fof(f2072, plain, $false, inference(avatar_sat_refutation, [], [f1632, f2064, f2067, f2071])).
fof(f2071, plain, ~ spl98_1, inference(avatar_contradiction_clause, [], [f2070])).
fof(f2070, plain, ($false | ~ spl98_1), inference(subsumption_resolution, [], [f2069, f674])).
fof(f674, plain, ~ (e2 = unit), inference(definition_unfolding, [], [f205, f239])).
fof(f239, plain, (e0 = unit), inference(cnf_transformation, [], [f3])).
fof(f3, plain, (e0 = unit), file('/home/ubuntu/library/tptp/Problems/ALG/ALG091+1.p', ax3)).
fof(f205, plain, ~ (e0 = e2), inference(cnf_transformation, [], [f1])).
fof(f1, plain, (~ (e3 = e4) & ~ (e2 = e4) & ~ (e2 = e3) & ~ (e1 = e4) & ~ (e1 = e3) & ~ (e1 = e2) & ~ (e0 = e4) & ~ (e0 = e3) & ~ (e0 = e2) & ~ (e0 = e1)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG091+1.p', ax1)).
fof(f2069, plain, ((e2 = unit) | ~ spl98_1), inference(backward_demodulation, [], [f676, f1623])).
fof(f1623, plain, ((e2 = op(e4, e4)) | ~ spl98_1), inference(avatar_component_clause, [], [f1621])).
fof(f1621, plain, (spl98_1 <=> (e2 = op(e4, e4))), introduced(avatar_definition, [new_symbols(naming, [spl98_1])])).
fof(f676, plain, (op(e4, e4) = unit), inference(definition_unfolding, [], [f238, f239])).
fof(f238, plain, (e0 = op(e4, e4)), inference(cnf_transformation, [], [f2])).
fof(f2, plain, ((e0 = op(e4, e4)) & (e2 = op(e4, e3)) & (e1 = op(e4, e2)) & (e3 = op(e4, e1)) & (e4 = op(e4, e0)) & (e1 = op(e3, e4)) & (e0 = op(e3, e3)) & (e4 = op(e3, e2)) & (e2 = op(e3, e1)) & (e3 = op(e3, e0)) & (e3 = op(e2, e4)) & (e1 = op(e2, e3)) & (e0 = op(e2, e2)) & (e4 = op(e2, e1)) & (e2 = op(e2, e0)) & (e2 = op(e1, e4)) & (e4 = op(e1, e3)) & (e3 = op(e1, e2)) & (e0 = op(e1, e1)) & (e1 = op(e1, e0)) & (e4 = op(e0, e4)) & (e3 = op(e0, e3)) & (e2 = op(e0, e2)) & (e1 = op(e0, e1)) & (e0 = op(e0, e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG091+1.p', ax2)).
fof(f2067, plain, ~ spl98_2, inference(avatar_contradiction_clause, [], [f2066])).
fof(f2066, plain, ($false | ~ spl98_2), inference(subsumption_resolution, [], [f2065, f673])).
fof(f673, plain, ~ (e3 = unit), inference(definition_unfolding, [], [f206, f239])).
fof(f206, plain, ~ (e0 = e3), inference(cnf_transformation, [], [f1])).
fof(f2065, plain, ((e3 = unit) | ~ spl98_2), inference(backward_demodulation, [], [f676, f1627])).
fof(f1627, plain, ((e3 = op(e4, e4)) | ~ spl98_2), inference(avatar_component_clause, [], [f1625])).
fof(f1625, plain, (spl98_2 <=> (e3 = op(e4, e4))), introduced(avatar_definition, [new_symbols(naming, [spl98_2])])).
fof(f2064, plain, (spl98_1 | spl98_2 | spl98_3), inference(avatar_split_clause, [], [f2063, f1629, f1625, f1621])).
fof(f1629, plain, (spl98_3 <=> (e4 = op(e4, e4))), introduced(avatar_definition, [new_symbols(naming, [spl98_3])])).
fof(f2063, plain, ((e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4))), inference(subsumption_resolution, [], [f2062, f963])).
fof(f963, plain, ~ sP76, inference(subsumption_resolution, [], [f705, f688])).
fof(f688, plain, (unit = op(unit, unit)), inference(definition_unfolding, [], [f214, f239, f239, f239])).
fof(f214, plain, (e0 = op(e0, e0)), inference(cnf_transformation, [], [f2])).
fof(f705, plain, (~ (unit = op(unit, unit)) | ~ sP76), inference(definition_unfolding, [], [f283, f239, f239, f239])).
fof(f283, plain, (~ (e0 = op(e0, e0)) | ~ sP76), inference(cnf_transformation, [], [f127])).
fof(f127, plain, ((~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP76), inference(nnf_transformation, [], [f83])).
fof(f83, plain, ((~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP76), inference(usedef, [], [e83])).
fof(e83, plain, (sP76 <=> (~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP76])])).
fof(f2062, plain, ((e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP76), inference(subsumption_resolution, [], [f2061, f962])).
fof(f962, plain, ~ sP77, inference(subsumption_resolution, [], [f961, f675])).
fof(f675, plain, ~ (e1 = unit), inference(definition_unfolding, [], [f204, f239])).
fof(f204, plain, ~ (e0 = e1), inference(cnf_transformation, [], [f1])).
fof(f961, plain, ((e1 = unit) | ~ sP77), inference(forward_demodulation, [], [f704, f688])).
fof(f704, plain, ((e1 = op(unit, unit)) | ~ sP77), inference(definition_unfolding, [], [f280, f239, f239])).
fof(f280, plain, ((e1 = op(e0, e0)) | ~ sP77), inference(cnf_transformation, [], [f126])).
fof(f126, plain, ((~ (e0 = op(e0, e1)) & (e1 = op(e0, e0))) | ~ sP77), inference(nnf_transformation, [], [f84])).
fof(f84, plain, ((~ (e0 = op(e0, e1)) & (e1 = op(e0, e0))) | ~ sP77), inference(usedef, [], [e84])).
fof(e84, plain, (sP77 <=> (~ (e0 = op(e0, e1)) & (e1 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP77])])).
fof(f2061, plain, ((e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP77 | sP76), inference(subsumption_resolution, [], [f2060, f959])).
fof(f959, plain, ~ sP78, inference(subsumption_resolution, [], [f958, f674])).
fof(f958, plain, ((e2 = unit) | ~ sP78), inference(forward_demodulation, [], [f702, f688])).
fof(f702, plain, ((e2 = op(unit, unit)) | ~ sP78), inference(definition_unfolding, [], [f278, f239, f239])).
fof(f278, plain, ((e2 = op(e0, e0)) | ~ sP78), inference(cnf_transformation, [], [f125])).
fof(f125, plain, ((~ (e0 = op(e0, e2)) & (e2 = op(e0, e0))) | ~ sP78), inference(nnf_transformation, [], [f85])).
fof(f85, plain, ((~ (e0 = op(e0, e2)) & (e2 = op(e0, e0))) | ~ sP78), inference(usedef, [], [e85])).
fof(e85, plain, (sP78 <=> (~ (e0 = op(e0, e2)) & (e2 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP78])])).
fof(f2060, plain, ((e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f2059, f956])).
fof(f956, plain, ~ sP79, inference(subsumption_resolution, [], [f955, f673])).
fof(f955, plain, ((e3 = unit) | ~ sP79), inference(forward_demodulation, [], [f700, f688])).
fof(f700, plain, ((e3 = op(unit, unit)) | ~ sP79), inference(definition_unfolding, [], [f276, f239, f239])).
fof(f276, plain, ((e3 = op(e0, e0)) | ~ sP79), inference(cnf_transformation, [], [f124])).
fof(f124, plain, ((~ (e0 = op(e0, e3)) & (e3 = op(e0, e0))) | ~ sP79), inference(nnf_transformation, [], [f86])).
fof(f86, plain, ((~ (e0 = op(e0, e3)) & (e3 = op(e0, e0))) | ~ sP79), inference(usedef, [], [e86])).
fof(e86, plain, (sP79 <=> (~ (e0 = op(e0, e3)) & (e3 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP79])])).
fof(f2059, plain, ((e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f2058, f953])).
fof(f953, plain, ~ sP80, inference(subsumption_resolution, [], [f952, f672])).
fof(f672, plain, ~ (e4 = unit), inference(definition_unfolding, [], [f207, f239])).
fof(f207, plain, ~ (e0 = e4), inference(cnf_transformation, [], [f1])).
fof(f952, plain, ((e4 = unit) | ~ sP80), inference(forward_demodulation, [], [f698, f688])).
fof(f698, plain, ((e4 = op(unit, unit)) | ~ sP80), inference(definition_unfolding, [], [f274, f239, f239])).
fof(f274, plain, ((e4 = op(e0, e0)) | ~ sP80), inference(cnf_transformation, [], [f123])).
fof(f123, plain, ((~ (e0 = op(e0, e4)) & (e4 = op(e0, e0))) | ~ sP80), inference(nnf_transformation, [], [f87])).
fof(f87, plain, ((~ (e0 = op(e0, e4)) & (e4 = op(e0, e0))) | ~ sP80), inference(usedef, [], [e87])).
fof(e87, plain, (sP80 <=> (~ (e0 = op(e0, e4)) & (e4 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP80])])).
fof(f2058, plain, ((e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f2057, f950])).
fof(f950, plain, ~ sP81, inference(subsumption_resolution, [], [f695, f683])).
fof(f683, plain, (e1 = op(e1, unit)), inference(definition_unfolding, [], [f219, f239])).
fof(f219, plain, (e1 = op(e1, e0)), inference(cnf_transformation, [], [f2])).
fof(f695, plain, (~ (e1 = op(e1, unit)) | ~ sP81), inference(definition_unfolding, [], [f273, f239])).
fof(f273, plain, (~ (e1 = op(e1, e0)) | ~ sP81), inference(cnf_transformation, [], [f122])).
fof(f122, plain, ((~ (e1 = op(e1, e0)) & (e0 = op(e1, e1))) | ~ sP81), inference(nnf_transformation, [], [f88])).
fof(f88, plain, ((~ (e1 = op(e1, e0)) & (e0 = op(e1, e1))) | ~ sP81), inference(usedef, [], [e88])).
fof(e88, plain, (sP81 <=> (~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP81])])).
fof(f2057, plain, ((e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f2056, f949])).
fof(f949, plain, ~ sP82, inference(subsumption_resolution, [], [f948, f675])).
fof(f948, plain, ((e1 = unit) | ~ sP82), inference(forward_demodulation, [], [f270, f682])).
fof(f682, plain, (op(e1, e1) = unit), inference(definition_unfolding, [], [f220, f239])).
fof(f220, plain, (e0 = op(e1, e1)), inference(cnf_transformation, [], [f2])).
fof(f270, plain, ((e1 = op(e1, e1)) | ~ sP82), inference(cnf_transformation, [], [f121])).
fof(f121, plain, ((~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP82), inference(nnf_transformation, [], [f89])).
fof(f89, plain, ((~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP82), inference(usedef, [], [e89])).
fof(e89, plain, (sP82 <=> (~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP82])])).
fof(f2056, plain, ((e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f2055, f946])).
fof(f946, plain, ~ sP83, inference(subsumption_resolution, [], [f945, f674])).
fof(f945, plain, ((e2 = unit) | ~ sP83), inference(forward_demodulation, [], [f268, f682])).
fof(f268, plain, ((e2 = op(e1, e1)) | ~ sP83), inference(cnf_transformation, [], [f120])).
fof(f120, plain, ((~ (e1 = op(e1, e2)) & (e2 = op(e1, e1))) | ~ sP83), inference(nnf_transformation, [], [f90])).
fof(f90, plain, ((~ (e1 = op(e1, e2)) & (e2 = op(e1, e1))) | ~ sP83), inference(usedef, [], [e90])).
fof(e90, plain, (sP83 <=> (~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP83])])).
fof(f2055, plain, ((e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f2054, f943])).
fof(f943, plain, ~ sP84, inference(subsumption_resolution, [], [f942, f673])).
fof(f942, plain, ((e3 = unit) | ~ sP84), inference(forward_demodulation, [], [f266, f682])).
fof(f266, plain, ((e3 = op(e1, e1)) | ~ sP84), inference(cnf_transformation, [], [f119])).
fof(f119, plain, ((~ (e1 = op(e1, e3)) & (e3 = op(e1, e1))) | ~ sP84), inference(nnf_transformation, [], [f91])).
fof(f91, plain, ((~ (e1 = op(e1, e3)) & (e3 = op(e1, e1))) | ~ sP84), inference(usedef, [], [e91])).
fof(e91, plain, (sP84 <=> (~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP84])])).
fof(f2054, plain, ((e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f2053, f940])).
fof(f940, plain, ~ sP85, inference(subsumption_resolution, [], [f939, f672])).
fof(f939, plain, ((e4 = unit) | ~ sP85), inference(forward_demodulation, [], [f264, f682])).
fof(f264, plain, ((e4 = op(e1, e1)) | ~ sP85), inference(cnf_transformation, [], [f118])).
fof(f118, plain, ((~ (e1 = op(e1, e4)) & (e4 = op(e1, e1))) | ~ sP85), inference(nnf_transformation, [], [f92])).
fof(f92, plain, ((~ (e1 = op(e1, e4)) & (e4 = op(e1, e1))) | ~ sP85), inference(usedef, [], [e92])).
fof(e92, plain, (sP85 <=> (~ (e1 = op(e1, e4)) & (e4 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP85])])).
fof(f2053, plain, ((e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f2052, f937])).
fof(f937, plain, ~ sP86, inference(subsumption_resolution, [], [f693, f681])).
fof(f681, plain, (e2 = op(e2, unit)), inference(definition_unfolding, [], [f224, f239])).
fof(f224, plain, (e2 = op(e2, e0)), inference(cnf_transformation, [], [f2])).
fof(f693, plain, (~ (e2 = op(e2, unit)) | ~ sP86), inference(definition_unfolding, [], [f263, f239])).
fof(f263, plain, (~ (e2 = op(e2, e0)) | ~ sP86), inference(cnf_transformation, [], [f117])).
fof(f117, plain, ((~ (e2 = op(e2, e0)) & (e0 = op(e2, e2))) | ~ sP86), inference(nnf_transformation, [], [f93])).
fof(f93, plain, ((~ (e2 = op(e2, e0)) & (e0 = op(e2, e2))) | ~ sP86), inference(usedef, [], [e93])).
fof(e93, plain, (sP86 <=> (~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP86])])).
fof(f2052, plain, ((e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f2051, f936])).
fof(f936, plain, ~ sP87, inference(subsumption_resolution, [], [f935, f675])).
fof(f935, plain, ((e1 = unit) | ~ sP87), inference(forward_demodulation, [], [f260, f680])).
fof(f680, plain, (op(e2, e2) = unit), inference(definition_unfolding, [], [f226, f239])).
fof(f226, plain, (e0 = op(e2, e2)), inference(cnf_transformation, [], [f2])).
fof(f260, plain, ((e1 = op(e2, e2)) | ~ sP87), inference(cnf_transformation, [], [f116])).
fof(f116, plain, ((~ (e2 = op(e2, e1)) & (e1 = op(e2, e2))) | ~ sP87), inference(nnf_transformation, [], [f94])).
fof(f94, plain, ((~ (e2 = op(e2, e1)) & (e1 = op(e2, e2))) | ~ sP87), inference(usedef, [], [e94])).
fof(e94, plain, (sP87 <=> (~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP87])])).
fof(f2051, plain, ((e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f2050, f933])).
fof(f933, plain, ~ sP88, inference(subsumption_resolution, [], [f932, f674])).
fof(f932, plain, ((e2 = unit) | ~ sP88), inference(forward_demodulation, [], [f258, f680])).
fof(f258, plain, ((e2 = op(e2, e2)) | ~ sP88), inference(cnf_transformation, [], [f115])).
fof(f115, plain, ((~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP88), inference(nnf_transformation, [], [f95])).
fof(f95, plain, ((~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP88), inference(usedef, [], [e95])).
fof(e95, plain, (sP88 <=> (~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP88])])).
fof(f2050, plain, ((e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f2049, f930])).
fof(f930, plain, ~ sP89, inference(subsumption_resolution, [], [f929, f673])).
fof(f929, plain, ((e3 = unit) | ~ sP89), inference(forward_demodulation, [], [f256, f680])).
fof(f256, plain, ((e3 = op(e2, e2)) | ~ sP89), inference(cnf_transformation, [], [f114])).
fof(f114, plain, ((~ (e2 = op(e2, e3)) & (e3 = op(e2, e2))) | ~ sP89), inference(nnf_transformation, [], [f96])).
fof(f96, plain, ((~ (e2 = op(e2, e3)) & (e3 = op(e2, e2))) | ~ sP89), inference(usedef, [], [e96])).
fof(e96, plain, (sP89 <=> (~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP89])])).
fof(f2049, plain, ((e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f2048, f927])).
fof(f927, plain, ~ sP90, inference(subsumption_resolution, [], [f926, f672])).
fof(f926, plain, ((e4 = unit) | ~ sP90), inference(forward_demodulation, [], [f254, f680])).
fof(f254, plain, ((e4 = op(e2, e2)) | ~ sP90), inference(cnf_transformation, [], [f113])).
fof(f113, plain, ((~ (e2 = op(e2, e4)) & (e4 = op(e2, e2))) | ~ sP90), inference(nnf_transformation, [], [f97])).
fof(f97, plain, ((~ (e2 = op(e2, e4)) & (e4 = op(e2, e2))) | ~ sP90), inference(usedef, [], [e97])).
fof(e97, plain, (sP90 <=> (~ (e2 = op(e2, e4)) & (e4 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP90])])).
fof(f2048, plain, ((e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f2047, f924])).
fof(f924, plain, ~ sP91, inference(subsumption_resolution, [], [f691, f679])).
fof(f679, plain, (e3 = op(e3, unit)), inference(definition_unfolding, [], [f229, f239])).
fof(f229, plain, (e3 = op(e3, e0)), inference(cnf_transformation, [], [f2])).
fof(f691, plain, (~ (e3 = op(e3, unit)) | ~ sP91), inference(definition_unfolding, [], [f253, f239])).
fof(f253, plain, (~ (e3 = op(e3, e0)) | ~ sP91), inference(cnf_transformation, [], [f112])).
fof(f112, plain, ((~ (e3 = op(e3, e0)) & (e0 = op(e3, e3))) | ~ sP91), inference(nnf_transformation, [], [f98])).
fof(f98, plain, ((~ (e3 = op(e3, e0)) & (e0 = op(e3, e3))) | ~ sP91), inference(usedef, [], [e98])).
fof(e98, plain, (sP91 <=> (~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP91])])).
fof(f2047, plain, ((e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f2046, f923])).
fof(f923, plain, ~ sP92, inference(subsumption_resolution, [], [f922, f675])).
fof(f922, plain, ((e1 = unit) | ~ sP92), inference(forward_demodulation, [], [f250, f678])).
fof(f678, plain, (op(e3, e3) = unit), inference(definition_unfolding, [], [f232, f239])).
fof(f232, plain, (e0 = op(e3, e3)), inference(cnf_transformation, [], [f2])).
fof(f250, plain, ((e1 = op(e3, e3)) | ~ sP92), inference(cnf_transformation, [], [f111])).
fof(f111, plain, ((~ (e3 = op(e3, e1)) & (e1 = op(e3, e3))) | ~ sP92), inference(nnf_transformation, [], [f99])).
fof(f99, plain, ((~ (e3 = op(e3, e1)) & (e1 = op(e3, e3))) | ~ sP92), inference(usedef, [], [e99])).
fof(e99, plain, (sP92 <=> (~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP92])])).
fof(f2046, plain, ((e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f2045, f920])).
fof(f920, plain, ~ sP93, inference(subsumption_resolution, [], [f919, f674])).
fof(f919, plain, ((e2 = unit) | ~ sP93), inference(forward_demodulation, [], [f248, f678])).
fof(f248, plain, ((e2 = op(e3, e3)) | ~ sP93), inference(cnf_transformation, [], [f110])).
fof(f110, plain, ((~ (e3 = op(e3, e2)) & (e2 = op(e3, e3))) | ~ sP93), inference(nnf_transformation, [], [f100])).
fof(f100, plain, ((~ (e3 = op(e3, e2)) & (e2 = op(e3, e3))) | ~ sP93), inference(usedef, [], [e100])).
fof(e100, plain, (sP93 <=> (~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP93])])).
fof(f2045, plain, ((e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f2044, f917])).
fof(f917, plain, ~ sP94, inference(subsumption_resolution, [], [f916, f673])).
fof(f916, plain, ((e3 = unit) | ~ sP94), inference(forward_demodulation, [], [f246, f678])).
fof(f246, plain, ((e3 = op(e3, e3)) | ~ sP94), inference(cnf_transformation, [], [f109])).
fof(f109, plain, ((~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ~ sP94), inference(nnf_transformation, [], [f101])).
fof(f101, plain, ((~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ~ sP94), inference(usedef, [], [e101])).
fof(e101, plain, (sP94 <=> (~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP94])])).
fof(f2044, plain, ((e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f2043, f914])).
fof(f914, plain, ~ sP95, inference(subsumption_resolution, [], [f913, f672])).
fof(f913, plain, ((e4 = unit) | ~ sP95), inference(forward_demodulation, [], [f244, f678])).
fof(f244, plain, ((e4 = op(e3, e3)) | ~ sP95), inference(cnf_transformation, [], [f108])).
fof(f108, plain, ((~ (e3 = op(e3, e4)) & (e4 = op(e3, e3))) | ~ sP95), inference(nnf_transformation, [], [f102])).
fof(f102, plain, ((~ (e3 = op(e3, e4)) & (e4 = op(e3, e3))) | ~ sP95), inference(usedef, [], [e102])).
fof(e102, plain, (sP95 <=> (~ (e3 = op(e3, e4)) & (e4 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP95])])).
fof(f2043, plain, ((e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f2042, f911])).
fof(f911, plain, ~ sP96, inference(subsumption_resolution, [], [f689, f677])).
fof(f677, plain, (e4 = op(e4, unit)), inference(definition_unfolding, [], [f234, f239])).
fof(f234, plain, (e4 = op(e4, e0)), inference(cnf_transformation, [], [f2])).
fof(f689, plain, (~ (e4 = op(e4, unit)) | ~ sP96), inference(definition_unfolding, [], [f243, f239])).
fof(f243, plain, (~ (e4 = op(e4, e0)) | ~ sP96), inference(cnf_transformation, [], [f107])).
fof(f107, plain, ((~ (e4 = op(e4, e0)) & (e0 = op(e4, e4))) | ~ sP96), inference(nnf_transformation, [], [f103])).
fof(f103, plain, ((~ (e4 = op(e4, e0)) & (e0 = op(e4, e4))) | ~ sP96), inference(usedef, [], [e103])).
fof(e103, plain, (sP96 <=> (~ (e4 = op(e4, e0)) & (e0 = op(e4, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP96])])).
fof(f2042, plain, ((e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f2041, f910])).
fof(f910, plain, ~ sP97, inference(subsumption_resolution, [], [f909, f675])).
fof(f909, plain, ((e1 = unit) | ~ sP97), inference(forward_demodulation, [], [f240, f676])).
fof(f240, plain, ((e1 = op(e4, e4)) | ~ sP97), inference(cnf_transformation, [], [f106])).
fof(f106, plain, ((~ (e4 = op(e4, e1)) & (e1 = op(e4, e4))) | ~ sP97), inference(nnf_transformation, [], [f104])).
fof(f104, plain, ((~ (e4 = op(e4, e1)) & (e1 = op(e4, e4))) | ~ sP97), inference(usedef, [], [e104])).
fof(e104, plain, (sP97 <=> (~ (e4 = op(e4, e1)) & (e1 = op(e4, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP97])])).
fof(f2041, plain, ((e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f2040, f1188])).
fof(f1188, plain, ~ sP0, inference(subsumption_resolution, [], [f890, f688])).
fof(f890, plain, (~ (unit = op(unit, unit)) | ~ sP0), inference(definition_unfolding, [], [f659, f239, f239, f239])).
fof(f659, plain, (~ (e0 = op(e0, e0)) | ~ sP0), inference(cnf_transformation, [], [f203])).
fof(f203, plain, ((~ (e4 = op(e0, e0)) & ~ (e3 = op(e0, e0)) & ~ (e2 = op(e0, e0)) & ~ (e1 = op(e0, e0)) & ~ (e0 = op(e0, e0))) | ~ sP0), inference(nnf_transformation, [], [f7])).
fof(f7, plain, ((~ (e4 = op(e0, e0)) & ~ (e3 = op(e0, e0)) & ~ (e2 = op(e0, e0)) & ~ (e1 = op(e0, e0)) & ~ (e0 = op(e0, e0))) | ~ sP0), inference(usedef, [], [e7])).
fof(e7, plain, (sP0 <=> (~ (e4 = op(e0, e0)) & ~ (e3 = op(e0, e0)) & ~ (e2 = op(e0, e0)) & ~ (e1 = op(e0, e0)) & ~ (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f2040, plain, (sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f2039, f1183])).
fof(f1183, plain, ~ sP1, inference(subsumption_resolution, [], [f884, f687])).
fof(f687, plain, (e1 = op(unit, e1)), inference(definition_unfolding, [], [f215, f239])).
fof(f215, plain, (e1 = op(e0, e1)), inference(cnf_transformation, [], [f2])).
fof(f884, plain, (~ (e1 = op(unit, e1)) | ~ sP1), inference(definition_unfolding, [], [f655, f239])).
fof(f655, plain, (~ (e1 = op(e0, e1)) | ~ sP1), inference(cnf_transformation, [], [f202])).
fof(f202, plain, ((~ (e4 = op(e0, e1)) & ~ (e3 = op(e0, e1)) & ~ (e2 = op(e0, e1)) & ~ (e1 = op(e0, e1)) & ~ (e0 = op(e0, e1))) | ~ sP1), inference(nnf_transformation, [], [f8])).
fof(f8, plain, ((~ (e4 = op(e0, e1)) & ~ (e3 = op(e0, e1)) & ~ (e2 = op(e0, e1)) & ~ (e1 = op(e0, e1)) & ~ (e0 = op(e0, e1))) | ~ sP1), inference(usedef, [], [e8])).
fof(e8, plain, (sP1 <=> (~ (e4 = op(e0, e1)) & ~ (e3 = op(e0, e1)) & ~ (e2 = op(e0, e1)) & ~ (e1 = op(e0, e1)) & ~ (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f2039, plain, (sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f2038, f1179])).
fof(f1179, plain, ~ sP2, inference(subsumption_resolution, [], [f878, f686])).
fof(f686, plain, (e2 = op(unit, e2)), inference(definition_unfolding, [], [f216, f239])).
fof(f216, plain, (e2 = op(e0, e2)), inference(cnf_transformation, [], [f2])).
fof(f878, plain, (~ (e2 = op(unit, e2)) | ~ sP2), inference(definition_unfolding, [], [f651, f239])).
fof(f651, plain, (~ (e2 = op(e0, e2)) | ~ sP2), inference(cnf_transformation, [], [f201])).
fof(f201, plain, ((~ (e4 = op(e0, e2)) & ~ (e3 = op(e0, e2)) & ~ (e2 = op(e0, e2)) & ~ (e1 = op(e0, e2)) & ~ (e0 = op(e0, e2))) | ~ sP2), inference(nnf_transformation, [], [f9])).
fof(f9, plain, ((~ (e4 = op(e0, e2)) & ~ (e3 = op(e0, e2)) & ~ (e2 = op(e0, e2)) & ~ (e1 = op(e0, e2)) & ~ (e0 = op(e0, e2))) | ~ sP2), inference(usedef, [], [e9])).
fof(e9, plain, (sP2 <=> (~ (e4 = op(e0, e2)) & ~ (e3 = op(e0, e2)) & ~ (e2 = op(e0, e2)) & ~ (e1 = op(e0, e2)) & ~ (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f2038, plain, (sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f2037, f1176])).
fof(f1176, plain, ~ sP3, inference(subsumption_resolution, [], [f872, f685])).
fof(f685, plain, (e3 = op(unit, e3)), inference(definition_unfolding, [], [f217, f239])).
fof(f217, plain, (e3 = op(e0, e3)), inference(cnf_transformation, [], [f2])).
fof(f872, plain, (~ (e3 = op(unit, e3)) | ~ sP3), inference(definition_unfolding, [], [f647, f239])).
fof(f647, plain, (~ (e3 = op(e0, e3)) | ~ sP3), inference(cnf_transformation, [], [f200])).
fof(f200, plain, ((~ (e4 = op(e0, e3)) & ~ (e3 = op(e0, e3)) & ~ (e2 = op(e0, e3)) & ~ (e1 = op(e0, e3)) & ~ (e0 = op(e0, e3))) | ~ sP3), inference(nnf_transformation, [], [f10])).
fof(f10, plain, ((~ (e4 = op(e0, e3)) & ~ (e3 = op(e0, e3)) & ~ (e2 = op(e0, e3)) & ~ (e1 = op(e0, e3)) & ~ (e0 = op(e0, e3))) | ~ sP3), inference(usedef, [], [e10])).
fof(e10, plain, (sP3 <=> (~ (e4 = op(e0, e3)) & ~ (e3 = op(e0, e3)) & ~ (e2 = op(e0, e3)) & ~ (e1 = op(e0, e3)) & ~ (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f2037, plain, (sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f2036, f1174])).
fof(f1174, plain, ~ sP4, inference(subsumption_resolution, [], [f866, f684])).
fof(f684, plain, (e4 = op(unit, e4)), inference(definition_unfolding, [], [f218, f239])).
fof(f218, plain, (e4 = op(e0, e4)), inference(cnf_transformation, [], [f2])).
fof(f866, plain, (~ (e4 = op(unit, e4)) | ~ sP4), inference(definition_unfolding, [], [f643, f239])).
fof(f643, plain, (~ (e4 = op(e0, e4)) | ~ sP4), inference(cnf_transformation, [], [f199])).
fof(f199, plain, ((~ (e4 = op(e0, e4)) & ~ (e3 = op(e0, e4)) & ~ (e2 = op(e0, e4)) & ~ (e1 = op(e0, e4)) & ~ (e0 = op(e0, e4))) | ~ sP4), inference(nnf_transformation, [], [f11])).
fof(f11, plain, ((~ (e4 = op(e0, e4)) & ~ (e3 = op(e0, e4)) & ~ (e2 = op(e0, e4)) & ~ (e1 = op(e0, e4)) & ~ (e0 = op(e0, e4))) | ~ sP4), inference(usedef, [], [e11])).
fof(e11, plain, (sP4 <=> (~ (e4 = op(e0, e4)) & ~ (e3 = op(e0, e4)) & ~ (e2 = op(e0, e4)) & ~ (e1 = op(e0, e4)) & ~ (e0 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f2036, plain, (sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f2035, f1173])).
fof(f1173, plain, ~ sP5, inference(subsumption_resolution, [], [f864, f683])).
fof(f864, plain, (~ (e1 = op(e1, unit)) | ~ sP5), inference(definition_unfolding, [], [f635, f239])).
fof(f635, plain, (~ (e1 = op(e1, e0)) | ~ sP5), inference(cnf_transformation, [], [f198])).
fof(f198, plain, ((~ (e4 = op(e1, e0)) & ~ (e3 = op(e1, e0)) & ~ (e2 = op(e1, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e1, e0))) | ~ sP5), inference(nnf_transformation, [], [f12])).
fof(f12, plain, ((~ (e4 = op(e1, e0)) & ~ (e3 = op(e1, e0)) & ~ (e2 = op(e1, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e1, e0))) | ~ sP5), inference(usedef, [], [e12])).
fof(e12, plain, (sP5 <=> (~ (e4 = op(e1, e0)) & ~ (e3 = op(e1, e0)) & ~ (e2 = op(e1, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f2035, plain, (sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f2034, f1169])).
fof(f1169, plain, ~ sP6, inference(subsumption_resolution, [], [f860, f682])).
fof(f860, plain, (~ (op(e1, e1) = unit) | ~ sP6), inference(definition_unfolding, [], [f629, f239])).
fof(f629, plain, (~ (e0 = op(e1, e1)) | ~ sP6), inference(cnf_transformation, [], [f197])).
fof(f197, plain, ((~ (e4 = op(e1, e1)) & ~ (e3 = op(e1, e1)) & ~ (e2 = op(e1, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e1, e1))) | ~ sP6), inference(nnf_transformation, [], [f13])).
fof(f13, plain, ((~ (e4 = op(e1, e1)) & ~ (e3 = op(e1, e1)) & ~ (e2 = op(e1, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e1, e1))) | ~ sP6), inference(usedef, [], [e13])).
fof(e13, plain, (sP6 <=> (~ (e4 = op(e1, e1)) & ~ (e3 = op(e1, e1)) & ~ (e2 = op(e1, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f2034, plain, (sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f2033, f1164])).
fof(f1164, plain, ~ sP7, inference(subsumption_resolution, [], [f627, f221])).
fof(f221, plain, (e3 = op(e1, e2)), inference(cnf_transformation, [], [f2])).
fof(f627, plain, (~ (e3 = op(e1, e2)) | ~ sP7), inference(cnf_transformation, [], [f196])).
fof(f196, plain, ((~ (e4 = op(e1, e2)) & ~ (e3 = op(e1, e2)) & ~ (e2 = op(e1, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e1, e2))) | ~ sP7), inference(nnf_transformation, [], [f14])).
fof(f14, plain, ((~ (e4 = op(e1, e2)) & ~ (e3 = op(e1, e2)) & ~ (e2 = op(e1, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e1, e2))) | ~ sP7), inference(usedef, [], [e14])).
fof(e14, plain, (sP7 <=> (~ (e4 = op(e1, e2)) & ~ (e3 = op(e1, e2)) & ~ (e2 = op(e1, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f2033, plain, (sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f2032, f1162])).
fof(f1162, plain, ~ sP8, inference(subsumption_resolution, [], [f623, f222])).
fof(f222, plain, (e4 = op(e1, e3)), inference(cnf_transformation, [], [f2])).
fof(f623, plain, (~ (e4 = op(e1, e3)) | ~ sP8), inference(cnf_transformation, [], [f195])).
fof(f195, plain, ((~ (e4 = op(e1, e3)) & ~ (e3 = op(e1, e3)) & ~ (e2 = op(e1, e3)) & ~ (e1 = op(e1, e3)) & ~ (e0 = op(e1, e3))) | ~ sP8), inference(nnf_transformation, [], [f15])).
fof(f15, plain, ((~ (e4 = op(e1, e3)) & ~ (e3 = op(e1, e3)) & ~ (e2 = op(e1, e3)) & ~ (e1 = op(e1, e3)) & ~ (e0 = op(e1, e3))) | ~ sP8), inference(usedef, [], [e15])).
fof(e15, plain, (sP8 <=> (~ (e4 = op(e1, e3)) & ~ (e3 = op(e1, e3)) & ~ (e2 = op(e1, e3)) & ~ (e1 = op(e1, e3)) & ~ (e0 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f2032, plain, (sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f2031, f1161])).
fof(f1161, plain, ~ sP9, inference(subsumption_resolution, [], [f616, f223])).
fof(f223, plain, (e2 = op(e1, e4)), inference(cnf_transformation, [], [f2])).
fof(f616, plain, (~ (e2 = op(e1, e4)) | ~ sP9), inference(cnf_transformation, [], [f194])).
fof(f194, plain, ((~ (e4 = op(e1, e4)) & ~ (e3 = op(e1, e4)) & ~ (e2 = op(e1, e4)) & ~ (e1 = op(e1, e4)) & ~ (e0 = op(e1, e4))) | ~ sP9), inference(nnf_transformation, [], [f16])).
fof(f16, plain, ((~ (e4 = op(e1, e4)) & ~ (e3 = op(e1, e4)) & ~ (e2 = op(e1, e4)) & ~ (e1 = op(e1, e4)) & ~ (e0 = op(e1, e4))) | ~ sP9), inference(usedef, [], [e16])).
fof(e16, plain, (sP9 <=> (~ (e4 = op(e1, e4)) & ~ (e3 = op(e1, e4)) & ~ (e2 = op(e1, e4)) & ~ (e1 = op(e1, e4)) & ~ (e0 = op(e1, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f2031, plain, (sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f2030, f1158])).
fof(f1158, plain, ~ sP10, inference(subsumption_resolution, [], [f854, f681])).
fof(f854, plain, (~ (e2 = op(e2, unit)) | ~ sP10), inference(definition_unfolding, [], [f611, f239])).
fof(f611, plain, (~ (e2 = op(e2, e0)) | ~ sP10), inference(cnf_transformation, [], [f193])).
fof(f193, plain, ((~ (e4 = op(e2, e0)) & ~ (e3 = op(e2, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e2, e0)) & ~ (e0 = op(e2, e0))) | ~ sP10), inference(nnf_transformation, [], [f17])).
fof(f17, plain, ((~ (e4 = op(e2, e0)) & ~ (e3 = op(e2, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e2, e0)) & ~ (e0 = op(e2, e0))) | ~ sP10), inference(usedef, [], [e17])).
fof(e17, plain, (sP10 <=> (~ (e4 = op(e2, e0)) & ~ (e3 = op(e2, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e2, e0)) & ~ (e0 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f2030, plain, (sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f2029, f1155])).
fof(f1155, plain, ~ sP11, inference(subsumption_resolution, [], [f608, f225])).
fof(f225, plain, (e4 = op(e2, e1)), inference(cnf_transformation, [], [f2])).
fof(f608, plain, (~ (e4 = op(e2, e1)) | ~ sP11), inference(cnf_transformation, [], [f192])).
fof(f192, plain, ((~ (e4 = op(e2, e1)) & ~ (e3 = op(e2, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e2, e1)) & ~ (e0 = op(e2, e1))) | ~ sP11), inference(nnf_transformation, [], [f18])).
fof(f18, plain, ((~ (e4 = op(e2, e1)) & ~ (e3 = op(e2, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e2, e1)) & ~ (e0 = op(e2, e1))) | ~ sP11), inference(usedef, [], [e18])).
fof(e18, plain, (sP11 <=> (~ (e4 = op(e2, e1)) & ~ (e3 = op(e2, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e2, e1)) & ~ (e0 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f2029, plain, (sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f2028, f1154])).
fof(f1154, plain, ~ sP12, inference(subsumption_resolution, [], [f850, f680])).
fof(f850, plain, (~ (op(e2, e2) = unit) | ~ sP12), inference(definition_unfolding, [], [f599, f239])).
fof(f599, plain, (~ (e0 = op(e2, e2)) | ~ sP12), inference(cnf_transformation, [], [f191])).
fof(f191, plain, ((~ (e4 = op(e2, e2)) & ~ (e3 = op(e2, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e2, e2)) & ~ (e0 = op(e2, e2))) | ~ sP12), inference(nnf_transformation, [], [f19])).
fof(f19, plain, ((~ (e4 = op(e2, e2)) & ~ (e3 = op(e2, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e2, e2)) & ~ (e0 = op(e2, e2))) | ~ sP12), inference(usedef, [], [e19])).
fof(e19, plain, (sP12 <=> (~ (e4 = op(e2, e2)) & ~ (e3 = op(e2, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e2, e2)) & ~ (e0 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f2028, plain, (sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f2027, f1149])).
fof(f1149, plain, ~ sP13, inference(subsumption_resolution, [], [f595, f227])).
fof(f227, plain, (e1 = op(e2, e3)), inference(cnf_transformation, [], [f2])).
fof(f595, plain, (~ (e1 = op(e2, e3)) | ~ sP13), inference(cnf_transformation, [], [f190])).
fof(f190, plain, ((~ (e4 = op(e2, e3)) & ~ (e3 = op(e2, e3)) & ~ (e2 = op(e2, e3)) & ~ (e1 = op(e2, e3)) & ~ (e0 = op(e2, e3))) | ~ sP13), inference(nnf_transformation, [], [f20])).
fof(f20, plain, ((~ (e4 = op(e2, e3)) & ~ (e3 = op(e2, e3)) & ~ (e2 = op(e2, e3)) & ~ (e1 = op(e2, e3)) & ~ (e0 = op(e2, e3))) | ~ sP13), inference(usedef, [], [e20])).
fof(e20, plain, (sP13 <=> (~ (e4 = op(e2, e3)) & ~ (e3 = op(e2, e3)) & ~ (e2 = op(e2, e3)) & ~ (e1 = op(e2, e3)) & ~ (e0 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f2027, plain, (sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f2026, f1145])).
fof(f1145, plain, ~ sP14, inference(subsumption_resolution, [], [f592, f228])).
fof(f228, plain, (e3 = op(e2, e4)), inference(cnf_transformation, [], [f2])).
fof(f592, plain, (~ (e3 = op(e2, e4)) | ~ sP14), inference(cnf_transformation, [], [f189])).
fof(f189, plain, ((~ (e4 = op(e2, e4)) & ~ (e3 = op(e2, e4)) & ~ (e2 = op(e2, e4)) & ~ (e1 = op(e2, e4)) & ~ (e0 = op(e2, e4))) | ~ sP14), inference(nnf_transformation, [], [f21])).
fof(f21, plain, ((~ (e4 = op(e2, e4)) & ~ (e3 = op(e2, e4)) & ~ (e2 = op(e2, e4)) & ~ (e1 = op(e2, e4)) & ~ (e0 = op(e2, e4))) | ~ sP14), inference(usedef, [], [e21])).
fof(e21, plain, (sP14 <=> (~ (e4 = op(e2, e4)) & ~ (e3 = op(e2, e4)) & ~ (e2 = op(e2, e4)) & ~ (e1 = op(e2, e4)) & ~ (e0 = op(e2, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f2026, plain, (sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f2025, f1143])).
fof(f1143, plain, ~ sP15, inference(subsumption_resolution, [], [f844, f679])).
fof(f844, plain, (~ (e3 = op(e3, unit)) | ~ sP15), inference(definition_unfolding, [], [f587, f239])).
fof(f587, plain, (~ (e3 = op(e3, e0)) | ~ sP15), inference(cnf_transformation, [], [f188])).
fof(f188, plain, ((~ (e4 = op(e3, e0)) & ~ (e3 = op(e3, e0)) & ~ (e2 = op(e3, e0)) & ~ (e1 = op(e3, e0)) & ~ (e0 = op(e3, e0))) | ~ sP15), inference(nnf_transformation, [], [f22])).
fof(f22, plain, ((~ (e4 = op(e3, e0)) & ~ (e3 = op(e3, e0)) & ~ (e2 = op(e3, e0)) & ~ (e1 = op(e3, e0)) & ~ (e0 = op(e3, e0))) | ~ sP15), inference(usedef, [], [e22])).
fof(e22, plain, (sP15 <=> (~ (e4 = op(e3, e0)) & ~ (e3 = op(e3, e0)) & ~ (e2 = op(e3, e0)) & ~ (e1 = op(e3, e0)) & ~ (e0 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f2025, plain, (sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f2024, f1141])).
fof(f1141, plain, ~ sP16, inference(subsumption_resolution, [], [f581, f230])).
fof(f230, plain, (e2 = op(e3, e1)), inference(cnf_transformation, [], [f2])).
fof(f581, plain, (~ (e2 = op(e3, e1)) | ~ sP16), inference(cnf_transformation, [], [f187])).
fof(f187, plain, ((~ (e4 = op(e3, e1)) & ~ (e3 = op(e3, e1)) & ~ (e2 = op(e3, e1)) & ~ (e1 = op(e3, e1)) & ~ (e0 = op(e3, e1))) | ~ sP16), inference(nnf_transformation, [], [f23])).
fof(f23, plain, ((~ (e4 = op(e3, e1)) & ~ (e3 = op(e3, e1)) & ~ (e2 = op(e3, e1)) & ~ (e1 = op(e3, e1)) & ~ (e0 = op(e3, e1))) | ~ sP16), inference(usedef, [], [e23])).
fof(e23, plain, (sP16 <=> (~ (e4 = op(e3, e1)) & ~ (e3 = op(e3, e1)) & ~ (e2 = op(e3, e1)) & ~ (e1 = op(e3, e1)) & ~ (e0 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f2024, plain, (sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f2023, f1138])).
fof(f1138, plain, ~ sP17, inference(subsumption_resolution, [], [f578, f231])).
fof(f231, plain, (e4 = op(e3, e2)), inference(cnf_transformation, [], [f2])).
fof(f578, plain, (~ (e4 = op(e3, e2)) | ~ sP17), inference(cnf_transformation, [], [f186])).
fof(f186, plain, ((~ (e4 = op(e3, e2)) & ~ (e3 = op(e3, e2)) & ~ (e2 = op(e3, e2)) & ~ (e1 = op(e3, e2)) & ~ (e0 = op(e3, e2))) | ~ sP17), inference(nnf_transformation, [], [f24])).
fof(f24, plain, ((~ (e4 = op(e3, e2)) & ~ (e3 = op(e3, e2)) & ~ (e2 = op(e3, e2)) & ~ (e1 = op(e3, e2)) & ~ (e0 = op(e3, e2))) | ~ sP17), inference(usedef, [], [e24])).
fof(e24, plain, (sP17 <=> (~ (e4 = op(e3, e2)) & ~ (e3 = op(e3, e2)) & ~ (e2 = op(e3, e2)) & ~ (e1 = op(e3, e2)) & ~ (e0 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f2023, plain, (sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f2022, f1137])).
fof(f1137, plain, ~ sP18, inference(subsumption_resolution, [], [f840, f678])).
fof(f840, plain, (~ (op(e3, e3) = unit) | ~ sP18), inference(definition_unfolding, [], [f569, f239])).
fof(f569, plain, (~ (e0 = op(e3, e3)) | ~ sP18), inference(cnf_transformation, [], [f185])).
fof(f185, plain, ((~ (e4 = op(e3, e3)) & ~ (e3 = op(e3, e3)) & ~ (e2 = op(e3, e3)) & ~ (e1 = op(e3, e3)) & ~ (e0 = op(e3, e3))) | ~ sP18), inference(nnf_transformation, [], [f25])).
fof(f25, plain, ((~ (e4 = op(e3, e3)) & ~ (e3 = op(e3, e3)) & ~ (e2 = op(e3, e3)) & ~ (e1 = op(e3, e3)) & ~ (e0 = op(e3, e3))) | ~ sP18), inference(usedef, [], [e25])).
fof(e25, plain, (sP18 <=> (~ (e4 = op(e3, e3)) & ~ (e3 = op(e3, e3)) & ~ (e2 = op(e3, e3)) & ~ (e1 = op(e3, e3)) & ~ (e0 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f2022, plain, (sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f2021, f1132])).
fof(f1132, plain, ~ sP19, inference(subsumption_resolution, [], [f565, f233])).
fof(f233, plain, (e1 = op(e3, e4)), inference(cnf_transformation, [], [f2])).
fof(f565, plain, (~ (e1 = op(e3, e4)) | ~ sP19), inference(cnf_transformation, [], [f184])).
fof(f184, plain, ((~ (e4 = op(e3, e4)) & ~ (e3 = op(e3, e4)) & ~ (e2 = op(e3, e4)) & ~ (e1 = op(e3, e4)) & ~ (e0 = op(e3, e4))) | ~ sP19), inference(nnf_transformation, [], [f26])).
fof(f26, plain, ((~ (e4 = op(e3, e4)) & ~ (e3 = op(e3, e4)) & ~ (e2 = op(e3, e4)) & ~ (e1 = op(e3, e4)) & ~ (e0 = op(e3, e4))) | ~ sP19), inference(usedef, [], [e26])).
fof(e26, plain, (sP19 <=> (~ (e4 = op(e3, e4)) & ~ (e3 = op(e3, e4)) & ~ (e2 = op(e3, e4)) & ~ (e1 = op(e3, e4)) & ~ (e0 = op(e3, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP19])])).
fof(f2021, plain, (sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f2020, f1128])).
fof(f1128, plain, ~ sP20, inference(subsumption_resolution, [], [f834, f677])).
fof(f834, plain, (~ (e4 = op(e4, unit)) | ~ sP20), inference(definition_unfolding, [], [f563, f239])).
fof(f563, plain, (~ (e4 = op(e4, e0)) | ~ sP20), inference(cnf_transformation, [], [f183])).
fof(f183, plain, ((~ (e4 = op(e4, e0)) & ~ (e3 = op(e4, e0)) & ~ (e2 = op(e4, e0)) & ~ (e1 = op(e4, e0)) & ~ (e0 = op(e4, e0))) | ~ sP20), inference(nnf_transformation, [], [f27])).
fof(f27, plain, ((~ (e4 = op(e4, e0)) & ~ (e3 = op(e4, e0)) & ~ (e2 = op(e4, e0)) & ~ (e1 = op(e4, e0)) & ~ (e0 = op(e4, e0))) | ~ sP20), inference(usedef, [], [e27])).
fof(e27, plain, (sP20 <=> (~ (e4 = op(e4, e0)) & ~ (e3 = op(e4, e0)) & ~ (e2 = op(e4, e0)) & ~ (e1 = op(e4, e0)) & ~ (e0 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP20])])).
fof(f2020, plain, (sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f2019, f1127])).
fof(f1127, plain, ~ sP21, inference(subsumption_resolution, [], [f557, f235])).
fof(f235, plain, (e3 = op(e4, e1)), inference(cnf_transformation, [], [f2])).
fof(f557, plain, (~ (e3 = op(e4, e1)) | ~ sP21), inference(cnf_transformation, [], [f182])).
fof(f182, plain, ((~ (e4 = op(e4, e1)) & ~ (e3 = op(e4, e1)) & ~ (e2 = op(e4, e1)) & ~ (e1 = op(e4, e1)) & ~ (e0 = op(e4, e1))) | ~ sP21), inference(nnf_transformation, [], [f28])).
fof(f28, plain, ((~ (e4 = op(e4, e1)) & ~ (e3 = op(e4, e1)) & ~ (e2 = op(e4, e1)) & ~ (e1 = op(e4, e1)) & ~ (e0 = op(e4, e1))) | ~ sP21), inference(usedef, [], [e28])).
fof(e28, plain, (sP21 <=> (~ (e4 = op(e4, e1)) & ~ (e3 = op(e4, e1)) & ~ (e2 = op(e4, e1)) & ~ (e1 = op(e4, e1)) & ~ (e0 = op(e4, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP21])])).
fof(f2019, plain, (sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f2018, f1125])).
fof(f1125, plain, ~ sP22, inference(subsumption_resolution, [], [f550, f236])).
fof(f236, plain, (e1 = op(e4, e2)), inference(cnf_transformation, [], [f2])).
fof(f550, plain, (~ (e1 = op(e4, e2)) | ~ sP22), inference(cnf_transformation, [], [f181])).
fof(f181, plain, ((~ (e4 = op(e4, e2)) & ~ (e3 = op(e4, e2)) & ~ (e2 = op(e4, e2)) & ~ (e1 = op(e4, e2)) & ~ (e0 = op(e4, e2))) | ~ sP22), inference(nnf_transformation, [], [f29])).
fof(f29, plain, ((~ (e4 = op(e4, e2)) & ~ (e3 = op(e4, e2)) & ~ (e2 = op(e4, e2)) & ~ (e1 = op(e4, e2)) & ~ (e0 = op(e4, e2))) | ~ sP22), inference(usedef, [], [e29])).
fof(e29, plain, (sP22 <=> (~ (e4 = op(e4, e2)) & ~ (e3 = op(e4, e2)) & ~ (e2 = op(e4, e2)) & ~ (e1 = op(e4, e2)) & ~ (e0 = op(e4, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP22])])).
fof(f2018, plain, (sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f2017, f1121])).
fof(f1121, plain, ~ sP23, inference(subsumption_resolution, [], [f546, f237])).
fof(f237, plain, (e2 = op(e4, e3)), inference(cnf_transformation, [], [f2])).
fof(f546, plain, (~ (e2 = op(e4, e3)) | ~ sP23), inference(cnf_transformation, [], [f180])).
fof(f180, plain, ((~ (e4 = op(e4, e3)) & ~ (e3 = op(e4, e3)) & ~ (e2 = op(e4, e3)) & ~ (e1 = op(e4, e3)) & ~ (e0 = op(e4, e3))) | ~ sP23), inference(nnf_transformation, [], [f30])).
fof(f30, plain, ((~ (e4 = op(e4, e3)) & ~ (e3 = op(e4, e3)) & ~ (e2 = op(e4, e3)) & ~ (e1 = op(e4, e3)) & ~ (e0 = op(e4, e3))) | ~ sP23), inference(usedef, [], [e30])).
fof(e30, plain, (sP23 <=> (~ (e4 = op(e4, e3)) & ~ (e3 = op(e4, e3)) & ~ (e2 = op(e4, e3)) & ~ (e1 = op(e4, e3)) & ~ (e0 = op(e4, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP23])])).
fof(f2017, plain, (sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f2016, f1118])).
fof(f1118, plain, ~ sP24, inference(subsumption_resolution, [], [f830, f676])).
fof(f830, plain, (~ (op(e4, e4) = unit) | ~ sP24), inference(definition_unfolding, [], [f539, f239])).
fof(f539, plain, (~ (e0 = op(e4, e4)) | ~ sP24), inference(cnf_transformation, [], [f179])).
fof(f179, plain, ((~ (e4 = op(e4, e4)) & ~ (e3 = op(e4, e4)) & ~ (e2 = op(e4, e4)) & ~ (e1 = op(e4, e4)) & ~ (e0 = op(e4, e4))) | ~ sP24), inference(nnf_transformation, [], [f31])).
fof(f31, plain, ((~ (e4 = op(e4, e4)) & ~ (e3 = op(e4, e4)) & ~ (e2 = op(e4, e4)) & ~ (e1 = op(e4, e4)) & ~ (e0 = op(e4, e4))) | ~ sP24), inference(usedef, [], [e31])).
fof(e31, plain, (sP24 <=> (~ (e4 = op(e4, e4)) & ~ (e3 = op(e4, e4)) & ~ (e2 = op(e4, e4)) & ~ (e1 = op(e4, e4)) & ~ (e0 = op(e4, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP24])])).
fof(f2016, plain, (sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f2015, f688])).
fof(f2015, plain, (~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f2014, f687])).
fof(f2014, plain, (~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f2013, f683])).
fof(f2013, plain, (~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f2012, f686])).
fof(f2012, plain, (~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f2011, f681])).
fof(f2011, plain, (~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f2010, f685])).
fof(f2010, plain, (~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f2009, f679])).
fof(f2009, plain, (~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f2008, f684])).
fof(f2008, plain, (~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f2007, f677])).
fof(f2007, plain, (~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f2006, f907])).
fof(f907, plain, ~ sP25, inference(trivial_inequality_removal, [], [f829])).
fof(f829, plain, (~ (unit = unit) | ~ sP25), inference(definition_unfolding, [], [f534, f239])).
fof(f534, plain, (~ (e0 = unit) | ~ sP25), inference(cnf_transformation, [], [f178])).
fof(f178, plain, ((~ (e4 = unit) & ~ (e3 = unit) & ~ (e2 = unit) & ~ (e1 = unit) & ~ (e0 = unit)) | ~ sP25), inference(nnf_transformation, [], [f32])).
fof(f32, plain, ((~ (e4 = unit) & ~ (e3 = unit) & ~ (e2 = unit) & ~ (e1 = unit) & ~ (e0 = unit)) | ~ sP25), inference(usedef, [], [e32])).
fof(e32, plain, (sP25 <=> (~ (e4 = unit) & ~ (e3 = unit) & ~ (e2 = unit) & ~ (e1 = unit) & ~ (e0 = unit))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP25])])).
fof(f2006, plain, (sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f2005, f1113])).
fof(f1113, plain, ~ sP26, inference(subsumption_resolution, [], [f828, f688])).
fof(f828, plain, (~ (unit = op(unit, unit)) | ~ sP26), inference(definition_unfolding, [], [f529, f239, f239, f239])).
fof(f529, plain, (~ (e0 = op(e0, e0)) | ~ sP26), inference(cnf_transformation, [], [f177])).
fof(f177, plain, ((~ (e0 = op(e0, e4)) & ~ (e0 = op(e0, e3)) & ~ (e0 = op(e0, e2)) & ~ (e0 = op(e0, e1)) & ~ (e0 = op(e0, e0))) | ~ sP26), inference(nnf_transformation, [], [f33])).
fof(f33, plain, ((~ (e0 = op(e0, e4)) & ~ (e0 = op(e0, e3)) & ~ (e0 = op(e0, e2)) & ~ (e0 = op(e0, e1)) & ~ (e0 = op(e0, e0))) | ~ sP26), inference(usedef, [], [e33])).
fof(e33, plain, (sP26 <=> (~ (e0 = op(e0, e4)) & ~ (e0 = op(e0, e3)) & ~ (e0 = op(e0, e2)) & ~ (e0 = op(e0, e1)) & ~ (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP26])])).
fof(f2005, plain, (sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f2004, f1108])).
fof(f1108, plain, ~ sP27, inference(subsumption_resolution, [], [f823, f688])).
fof(f823, plain, (~ (unit = op(unit, unit)) | ~ sP27), inference(definition_unfolding, [], [f524, f239, f239, f239])).
fof(f524, plain, (~ (e0 = op(e0, e0)) | ~ sP27), inference(cnf_transformation, [], [f176])).
fof(f176, plain, ((~ (e0 = op(e4, e0)) & ~ (e0 = op(e3, e0)) & ~ (e0 = op(e2, e0)) & ~ (e0 = op(e1, e0)) & ~ (e0 = op(e0, e0))) | ~ sP27), inference(nnf_transformation, [], [f34])).
fof(f34, plain, ((~ (e0 = op(e4, e0)) & ~ (e0 = op(e3, e0)) & ~ (e0 = op(e2, e0)) & ~ (e0 = op(e1, e0)) & ~ (e0 = op(e0, e0))) | ~ sP27), inference(usedef, [], [e34])).
fof(e34, plain, (sP27 <=> (~ (e0 = op(e4, e0)) & ~ (e0 = op(e3, e0)) & ~ (e0 = op(e2, e0)) & ~ (e0 = op(e1, e0)) & ~ (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP27])])).
fof(f2004, plain, (sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f2003, f1103])).
fof(f1103, plain, ~ sP28, inference(subsumption_resolution, [], [f817, f687])).
fof(f817, plain, (~ (e1 = op(unit, e1)) | ~ sP28), inference(definition_unfolding, [], [f520, f239])).
fof(f520, plain, (~ (e1 = op(e0, e1)) | ~ sP28), inference(cnf_transformation, [], [f175])).
fof(f175, plain, ((~ (e1 = op(e0, e4)) & ~ (e1 = op(e0, e3)) & ~ (e1 = op(e0, e2)) & ~ (e1 = op(e0, e1)) & ~ (e1 = op(e0, e0))) | ~ sP28), inference(nnf_transformation, [], [f35])).
fof(f35, plain, ((~ (e1 = op(e0, e4)) & ~ (e1 = op(e0, e3)) & ~ (e1 = op(e0, e2)) & ~ (e1 = op(e0, e1)) & ~ (e1 = op(e0, e0))) | ~ sP28), inference(usedef, [], [e35])).
fof(e35, plain, (sP28 <=> (~ (e1 = op(e0, e4)) & ~ (e1 = op(e0, e3)) & ~ (e1 = op(e0, e2)) & ~ (e1 = op(e0, e1)) & ~ (e1 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP28])])).
fof(f2003, plain, (sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f2002, f1099])).
fof(f1099, plain, ~ sP29, inference(subsumption_resolution, [], [f812, f683])).
fof(f812, plain, (~ (e1 = op(e1, unit)) | ~ sP29), inference(definition_unfolding, [], [f515, f239])).
fof(f515, plain, (~ (e1 = op(e1, e0)) | ~ sP29), inference(cnf_transformation, [], [f174])).
fof(f174, plain, ((~ (e1 = op(e4, e0)) & ~ (e1 = op(e3, e0)) & ~ (e1 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e1 = op(e0, e0))) | ~ sP29), inference(nnf_transformation, [], [f36])).
fof(f36, plain, ((~ (e1 = op(e4, e0)) & ~ (e1 = op(e3, e0)) & ~ (e1 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e1 = op(e0, e0))) | ~ sP29), inference(usedef, [], [e36])).
fof(e36, plain, (sP29 <=> (~ (e1 = op(e4, e0)) & ~ (e1 = op(e3, e0)) & ~ (e1 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e1 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP29])])).
fof(f2002, plain, (sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f2001, f1095])).
fof(f1095, plain, ~ sP30, inference(subsumption_resolution, [], [f806, f686])).
fof(f806, plain, (~ (e2 = op(unit, e2)) | ~ sP30), inference(definition_unfolding, [], [f511, f239])).
fof(f511, plain, (~ (e2 = op(e0, e2)) | ~ sP30), inference(cnf_transformation, [], [f173])).
fof(f173, plain, ((~ (e2 = op(e0, e4)) & ~ (e2 = op(e0, e3)) & ~ (e2 = op(e0, e2)) & ~ (e2 = op(e0, e1)) & ~ (e2 = op(e0, e0))) | ~ sP30), inference(nnf_transformation, [], [f37])).
fof(f37, plain, ((~ (e2 = op(e0, e4)) & ~ (e2 = op(e0, e3)) & ~ (e2 = op(e0, e2)) & ~ (e2 = op(e0, e1)) & ~ (e2 = op(e0, e0))) | ~ sP30), inference(usedef, [], [e37])).
fof(e37, plain, (sP30 <=> (~ (e2 = op(e0, e4)) & ~ (e2 = op(e0, e3)) & ~ (e2 = op(e0, e2)) & ~ (e2 = op(e0, e1)) & ~ (e2 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP30])])).
fof(f2001, plain, (sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f2000, f1092])).
fof(f1092, plain, ~ sP31, inference(subsumption_resolution, [], [f801, f681])).
fof(f801, plain, (~ (e2 = op(e2, unit)) | ~ sP31), inference(definition_unfolding, [], [f506, f239])).
fof(f506, plain, (~ (e2 = op(e2, e0)) | ~ sP31), inference(cnf_transformation, [], [f172])).
fof(f172, plain, ((~ (e2 = op(e4, e0)) & ~ (e2 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e2 = op(e1, e0)) & ~ (e2 = op(e0, e0))) | ~ sP31), inference(nnf_transformation, [], [f38])).
fof(f38, plain, ((~ (e2 = op(e4, e0)) & ~ (e2 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e2 = op(e1, e0)) & ~ (e2 = op(e0, e0))) | ~ sP31), inference(usedef, [], [e38])).
fof(e38, plain, (sP31 <=> (~ (e2 = op(e4, e0)) & ~ (e2 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e2 = op(e1, e0)) & ~ (e2 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP31])])).
fof(f2000, plain, (sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1999, f1089])).
fof(f1089, plain, ~ sP32, inference(subsumption_resolution, [], [f795, f685])).
fof(f795, plain, (~ (e3 = op(unit, e3)) | ~ sP32), inference(definition_unfolding, [], [f502, f239])).
fof(f502, plain, (~ (e3 = op(e0, e3)) | ~ sP32), inference(cnf_transformation, [], [f171])).
fof(f171, plain, ((~ (e3 = op(e0, e4)) & ~ (e3 = op(e0, e3)) & ~ (e3 = op(e0, e2)) & ~ (e3 = op(e0, e1)) & ~ (e3 = op(e0, e0))) | ~ sP32), inference(nnf_transformation, [], [f39])).
fof(f39, plain, ((~ (e3 = op(e0, e4)) & ~ (e3 = op(e0, e3)) & ~ (e3 = op(e0, e2)) & ~ (e3 = op(e0, e1)) & ~ (e3 = op(e0, e0))) | ~ sP32), inference(usedef, [], [e39])).
fof(e39, plain, (sP32 <=> (~ (e3 = op(e0, e4)) & ~ (e3 = op(e0, e3)) & ~ (e3 = op(e0, e2)) & ~ (e3 = op(e0, e1)) & ~ (e3 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP32])])).
fof(f1999, plain, (sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1998, f1087])).
fof(f1087, plain, ~ sP33, inference(subsumption_resolution, [], [f790, f679])).
fof(f790, plain, (~ (e3 = op(e3, unit)) | ~ sP33), inference(definition_unfolding, [], [f497, f239])).
fof(f497, plain, (~ (e3 = op(e3, e0)) | ~ sP33), inference(cnf_transformation, [], [f170])).
fof(f170, plain, ((~ (e3 = op(e4, e0)) & ~ (e3 = op(e3, e0)) & ~ (e3 = op(e2, e0)) & ~ (e3 = op(e1, e0)) & ~ (e3 = op(e0, e0))) | ~ sP33), inference(nnf_transformation, [], [f40])).
fof(f40, plain, ((~ (e3 = op(e4, e0)) & ~ (e3 = op(e3, e0)) & ~ (e3 = op(e2, e0)) & ~ (e3 = op(e1, e0)) & ~ (e3 = op(e0, e0))) | ~ sP33), inference(usedef, [], [e40])).
fof(e40, plain, (sP33 <=> (~ (e3 = op(e4, e0)) & ~ (e3 = op(e3, e0)) & ~ (e3 = op(e2, e0)) & ~ (e3 = op(e1, e0)) & ~ (e3 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP33])])).
fof(f1998, plain, (sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1997, f1085])).
fof(f1085, plain, ~ sP34, inference(subsumption_resolution, [], [f784, f684])).
fof(f784, plain, (~ (e4 = op(unit, e4)) | ~ sP34), inference(definition_unfolding, [], [f493, f239])).
fof(f493, plain, (~ (e4 = op(e0, e4)) | ~ sP34), inference(cnf_transformation, [], [f169])).
fof(f169, plain, ((~ (e4 = op(e0, e4)) & ~ (e4 = op(e0, e3)) & ~ (e4 = op(e0, e2)) & ~ (e4 = op(e0, e1)) & ~ (e4 = op(e0, e0))) | ~ sP34), inference(nnf_transformation, [], [f41])).
fof(f41, plain, ((~ (e4 = op(e0, e4)) & ~ (e4 = op(e0, e3)) & ~ (e4 = op(e0, e2)) & ~ (e4 = op(e0, e1)) & ~ (e4 = op(e0, e0))) | ~ sP34), inference(usedef, [], [e41])).
fof(e41, plain, (sP34 <=> (~ (e4 = op(e0, e4)) & ~ (e4 = op(e0, e3)) & ~ (e4 = op(e0, e2)) & ~ (e4 = op(e0, e1)) & ~ (e4 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP34])])).
fof(f1997, plain, (sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1996, f1084])).
fof(f1084, plain, ~ sP35, inference(subsumption_resolution, [], [f779, f677])).
fof(f779, plain, (~ (e4 = op(e4, unit)) | ~ sP35), inference(definition_unfolding, [], [f488, f239])).
fof(f488, plain, (~ (e4 = op(e4, e0)) | ~ sP35), inference(cnf_transformation, [], [f168])).
fof(f168, plain, ((~ (e4 = op(e4, e0)) & ~ (e4 = op(e3, e0)) & ~ (e4 = op(e2, e0)) & ~ (e4 = op(e1, e0)) & ~ (e4 = op(e0, e0))) | ~ sP35), inference(nnf_transformation, [], [f42])).
fof(f42, plain, ((~ (e4 = op(e4, e0)) & ~ (e4 = op(e3, e0)) & ~ (e4 = op(e2, e0)) & ~ (e4 = op(e1, e0)) & ~ (e4 = op(e0, e0))) | ~ sP35), inference(usedef, [], [e42])).
fof(e42, plain, (sP35 <=> (~ (e4 = op(e4, e0)) & ~ (e4 = op(e3, e0)) & ~ (e4 = op(e2, e0)) & ~ (e4 = op(e1, e0)) & ~ (e4 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP35])])).
fof(f1996, plain, (sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1995, f1083])).
fof(f1083, plain, ~ sP36, inference(subsumption_resolution, [], [f777, f682])).
fof(f777, plain, (~ (op(e1, e1) = unit) | ~ sP36), inference(definition_unfolding, [], [f480, f239])).
fof(f480, plain, (~ (e0 = op(e1, e1)) | ~ sP36), inference(cnf_transformation, [], [f167])).
fof(f167, plain, ((~ (e0 = op(e1, e4)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e1, e0))) | ~ sP36), inference(nnf_transformation, [], [f43])).
fof(f43, plain, ((~ (e0 = op(e1, e4)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e1, e0))) | ~ sP36), inference(usedef, [], [e43])).
fof(e43, plain, (sP36 <=> (~ (e0 = op(e1, e4)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP36])])).
fof(f1995, plain, (sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1994, f1079])).
fof(f1079, plain, ~ sP37, inference(subsumption_resolution, [], [f772, f682])).
fof(f772, plain, (~ (op(e1, e1) = unit) | ~ sP37), inference(definition_unfolding, [], [f475, f239])).
fof(f475, plain, (~ (e0 = op(e1, e1)) | ~ sP37), inference(cnf_transformation, [], [f166])).
fof(f166, plain, ((~ (e0 = op(e4, e1)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e0, e1))) | ~ sP37), inference(nnf_transformation, [], [f44])).
fof(f44, plain, ((~ (e0 = op(e4, e1)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e0, e1))) | ~ sP37), inference(usedef, [], [e44])).
fof(e44, plain, (sP37 <=> (~ (e0 = op(e4, e1)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP37])])).
fof(f1994, plain, (sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1993, f1075])).
fof(f1075, plain, ~ sP38, inference(subsumption_resolution, [], [f768, f683])).
fof(f768, plain, (~ (e1 = op(e1, unit)) | ~ sP38), inference(definition_unfolding, [], [f469, f239])).
fof(f469, plain, (~ (e1 = op(e1, e0)) | ~ sP38), inference(cnf_transformation, [], [f165])).
fof(f165, plain, ((~ (e1 = op(e1, e4)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e1, e0))) | ~ sP38), inference(nnf_transformation, [], [f45])).
fof(f45, plain, ((~ (e1 = op(e1, e4)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e1, e0))) | ~ sP38), inference(usedef, [], [e45])).
fof(e45, plain, (sP38 <=> (~ (e1 = op(e1, e4)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP38])])).
fof(f1993, plain, (sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1992, f1070])).
fof(f1070, plain, ~ sP39, inference(subsumption_resolution, [], [f767, f687])).
fof(f767, plain, (~ (e1 = op(unit, e1)) | ~ sP39), inference(definition_unfolding, [], [f464, f239])).
fof(f464, plain, (~ (e1 = op(e0, e1)) | ~ sP39), inference(cnf_transformation, [], [f164])).
fof(f164, plain, ((~ (e1 = op(e4, e1)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e0, e1))) | ~ sP39), inference(nnf_transformation, [], [f46])).
fof(f46, plain, ((~ (e1 = op(e4, e1)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e0, e1))) | ~ sP39), inference(usedef, [], [e46])).
fof(e46, plain, (sP39 <=> (~ (e1 = op(e4, e1)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP39])])).
fof(f1992, plain, (sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1991, f1065])).
fof(f1065, plain, ~ sP40, inference(subsumption_resolution, [], [f463, f223])).
fof(f463, plain, (~ (e2 = op(e1, e4)) | ~ sP40), inference(cnf_transformation, [], [f163])).
fof(f163, plain, ((~ (e2 = op(e1, e4)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e1, e0))) | ~ sP40), inference(nnf_transformation, [], [f47])).
fof(f47, plain, ((~ (e2 = op(e1, e4)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e1, e0))) | ~ sP40), inference(usedef, [], [e47])).
fof(e47, plain, (sP40 <=> (~ (e2 = op(e1, e4)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP40])])).
fof(f1991, plain, (sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1990, f1064])).
fof(f1064, plain, ~ sP41, inference(subsumption_resolution, [], [f457, f230])).
fof(f457, plain, (~ (e2 = op(e3, e1)) | ~ sP41), inference(cnf_transformation, [], [f162])).
fof(f162, plain, ((~ (e2 = op(e4, e1)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e0, e1))) | ~ sP41), inference(nnf_transformation, [], [f48])).
fof(f48, plain, ((~ (e2 = op(e4, e1)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e0, e1))) | ~ sP41), inference(usedef, [], [e48])).
fof(e48, plain, (sP41 <=> (~ (e2 = op(e4, e1)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP41])])).
fof(f1990, plain, (sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1989, f1062])).
fof(f1062, plain, ~ sP42, inference(subsumption_resolution, [], [f451, f221])).
fof(f451, plain, (~ (e3 = op(e1, e2)) | ~ sP42), inference(cnf_transformation, [], [f161])).
fof(f161, plain, ((~ (e3 = op(e1, e4)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e1, e0))) | ~ sP42), inference(nnf_transformation, [], [f49])).
fof(f49, plain, ((~ (e3 = op(e1, e4)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e1, e0))) | ~ sP42), inference(usedef, [], [e49])).
fof(e49, plain, (sP42 <=> (~ (e3 = op(e1, e4)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP42])])).
fof(f1989, plain, (sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1988, f1059])).
fof(f1059, plain, ~ sP43, inference(subsumption_resolution, [], [f448, f235])).
fof(f448, plain, (~ (e3 = op(e4, e1)) | ~ sP43), inference(cnf_transformation, [], [f160])).
fof(f160, plain, ((~ (e3 = op(e4, e1)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e0, e1))) | ~ sP43), inference(nnf_transformation, [], [f50])).
fof(f50, plain, ((~ (e3 = op(e4, e1)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e0, e1))) | ~ sP43), inference(usedef, [], [e50])).
fof(e50, plain, (sP43 <=> (~ (e3 = op(e4, e1)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP43])])).
fof(f1988, plain, (sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1987, f1058])).
fof(f1058, plain, ~ sP44, inference(subsumption_resolution, [], [f442, f222])).
fof(f442, plain, (~ (e4 = op(e1, e3)) | ~ sP44), inference(cnf_transformation, [], [f159])).
fof(f159, plain, ((~ (e4 = op(e1, e4)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e1, e0))) | ~ sP44), inference(nnf_transformation, [], [f51])).
fof(f51, plain, ((~ (e4 = op(e1, e4)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e1, e0))) | ~ sP44), inference(usedef, [], [e51])).
fof(e51, plain, (sP44 <=> (~ (e4 = op(e1, e4)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP44])])).
fof(f1987, plain, (sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1986, f1056])).
fof(f1056, plain, ~ sP45, inference(subsumption_resolution, [], [f436, f225])).
fof(f436, plain, (~ (e4 = op(e2, e1)) | ~ sP45), inference(cnf_transformation, [], [f158])).
fof(f158, plain, ((~ (e4 = op(e4, e1)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e0, e1))) | ~ sP45), inference(nnf_transformation, [], [f52])).
fof(f52, plain, ((~ (e4 = op(e4, e1)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e0, e1))) | ~ sP45), inference(usedef, [], [e52])).
fof(e52, plain, (sP45 <=> (~ (e4 = op(e4, e1)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP45])])).
fof(f1986, plain, (sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1985, f1053])).
fof(f1053, plain, ~ sP46, inference(subsumption_resolution, [], [f758, f680])).
fof(f758, plain, (~ (op(e2, e2) = unit) | ~ sP46), inference(definition_unfolding, [], [f431, f239])).
fof(f431, plain, (~ (e0 = op(e2, e2)) | ~ sP46), inference(cnf_transformation, [], [f157])).
fof(f157, plain, ((~ (e0 = op(e2, e4)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e2, e0))) | ~ sP46), inference(nnf_transformation, [], [f53])).
fof(f53, plain, ((~ (e0 = op(e2, e4)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e2, e0))) | ~ sP46), inference(usedef, [], [e53])).
fof(e53, plain, (sP46 <=> (~ (e0 = op(e2, e4)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP46])])).
fof(f1985, plain, (sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1984, f1050])).
fof(f1050, plain, ~ sP47, inference(subsumption_resolution, [], [f753, f680])).
fof(f753, plain, (~ (op(e2, e2) = unit) | ~ sP47), inference(definition_unfolding, [], [f426, f239])).
fof(f426, plain, (~ (e0 = op(e2, e2)) | ~ sP47), inference(cnf_transformation, [], [f156])).
fof(f156, plain, ((~ (e0 = op(e4, e2)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e0, e2))) | ~ sP47), inference(nnf_transformation, [], [f54])).
fof(f54, plain, ((~ (e0 = op(e4, e2)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e0, e2))) | ~ sP47), inference(usedef, [], [e54])).
fof(e54, plain, (sP47 <=> (~ (e0 = op(e4, e2)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP47])])).
fof(f1984, plain, (sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1983, f1047])).
fof(f1047, plain, ~ sP48, inference(subsumption_resolution, [], [f422, f227])).
fof(f422, plain, (~ (e1 = op(e2, e3)) | ~ sP48), inference(cnf_transformation, [], [f155])).
fof(f155, plain, ((~ (e1 = op(e2, e4)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e2, e0))) | ~ sP48), inference(nnf_transformation, [], [f55])).
fof(f55, plain, ((~ (e1 = op(e2, e4)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e2, e0))) | ~ sP48), inference(usedef, [], [e55])).
fof(e55, plain, (sP48 <=> (~ (e1 = op(e2, e4)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP48])])).
fof(f1983, plain, (sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1982, f1045])).
fof(f1045, plain, ~ sP49, inference(subsumption_resolution, [], [f418, f236])).
fof(f418, plain, (~ (e1 = op(e4, e2)) | ~ sP49), inference(cnf_transformation, [], [f154])).
fof(f154, plain, ((~ (e1 = op(e4, e2)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e0, e2))) | ~ sP49), inference(nnf_transformation, [], [f56])).
fof(f56, plain, ((~ (e1 = op(e4, e2)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e0, e2))) | ~ sP49), inference(usedef, [], [e56])).
fof(e56, plain, (sP49 <=> (~ (e1 = op(e4, e2)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP49])])).
fof(f1982, plain, (sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1981, f1044])).
fof(f1044, plain, ~ sP50, inference(subsumption_resolution, [], [f748, f681])).
fof(f748, plain, (~ (e2 = op(e2, unit)) | ~ sP50), inference(definition_unfolding, [], [f409, f239])).
fof(f409, plain, (~ (e2 = op(e2, e0)) | ~ sP50), inference(cnf_transformation, [], [f153])).
fof(f153, plain, ((~ (e2 = op(e2, e4)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e2, e0))) | ~ sP50), inference(nnf_transformation, [], [f57])).
fof(f57, plain, ((~ (e2 = op(e2, e4)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e2, e0))) | ~ sP50), inference(usedef, [], [e57])).
fof(e57, plain, (sP50 <=> (~ (e2 = op(e2, e4)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP50])])).
fof(f1981, plain, (sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1980, f1039])).
fof(f1039, plain, ~ sP51, inference(subsumption_resolution, [], [f747, f686])).
fof(f747, plain, (~ (e2 = op(unit, e2)) | ~ sP51), inference(definition_unfolding, [], [f404, f239])).
fof(f404, plain, (~ (e2 = op(e0, e2)) | ~ sP51), inference(cnf_transformation, [], [f152])).
fof(f152, plain, ((~ (e2 = op(e4, e2)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e0, e2))) | ~ sP51), inference(nnf_transformation, [], [f58])).
fof(f58, plain, ((~ (e2 = op(e4, e2)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e0, e2))) | ~ sP51), inference(usedef, [], [e58])).
fof(e58, plain, (sP51 <=> (~ (e2 = op(e4, e2)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP51])])).
fof(f1980, plain, (sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1979, f1034])).
fof(f1034, plain, ~ sP52, inference(subsumption_resolution, [], [f403, f228])).
fof(f403, plain, (~ (e3 = op(e2, e4)) | ~ sP52), inference(cnf_transformation, [], [f151])).
fof(f151, plain, ((~ (e3 = op(e2, e4)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e2, e0))) | ~ sP52), inference(nnf_transformation, [], [f59])).
fof(f59, plain, ((~ (e3 = op(e2, e4)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e2, e0))) | ~ sP52), inference(usedef, [], [e59])).
fof(e59, plain, (sP52 <=> (~ (e3 = op(e2, e4)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP52])])).
fof(f1979, plain, (sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1978, f1033])).
fof(f1033, plain, ~ sP53, inference(subsumption_resolution, [], [f395, f221])).
fof(f395, plain, (~ (e3 = op(e1, e2)) | ~ sP53), inference(cnf_transformation, [], [f150])).
fof(f150, plain, ((~ (e3 = op(e4, e2)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e0, e2))) | ~ sP53), inference(nnf_transformation, [], [f60])).
fof(f60, plain, ((~ (e3 = op(e4, e2)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e0, e2))) | ~ sP53), inference(usedef, [], [e60])).
fof(e60, plain, (sP53 <=> (~ (e3 = op(e4, e2)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP53])])).
fof(f1978, plain, (sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1977, f1029])).
fof(f1029, plain, ~ sP54, inference(subsumption_resolution, [], [f390, f225])).
fof(f390, plain, (~ (e4 = op(e2, e1)) | ~ sP54), inference(cnf_transformation, [], [f149])).
fof(f149, plain, ((~ (e4 = op(e2, e4)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e2, e0))) | ~ sP54), inference(nnf_transformation, [], [f61])).
fof(f61, plain, ((~ (e4 = op(e2, e4)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e2, e0))) | ~ sP54), inference(usedef, [], [e61])).
fof(e61, plain, (sP54 <=> (~ (e4 = op(e2, e4)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP54])])).
fof(f1977, plain, (sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1976, f1025])).
fof(f1025, plain, ~ sP55, inference(subsumption_resolution, [], [f387, f231])).
fof(f387, plain, (~ (e4 = op(e3, e2)) | ~ sP55), inference(cnf_transformation, [], [f148])).
fof(f148, plain, ((~ (e4 = op(e4, e2)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e0, e2))) | ~ sP55), inference(nnf_transformation, [], [f62])).
fof(f62, plain, ((~ (e4 = op(e4, e2)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e0, e2))) | ~ sP55), inference(usedef, [], [e62])).
fof(e62, plain, (sP55 <=> (~ (e4 = op(e4, e2)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP55])])).
fof(f1976, plain, (sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1975, f1023])).
fof(f1023, plain, ~ sP56, inference(subsumption_resolution, [], [f739, f678])).
fof(f739, plain, (~ (op(e3, e3) = unit) | ~ sP56), inference(definition_unfolding, [], [f382, f239])).
fof(f382, plain, (~ (e0 = op(e3, e3)) | ~ sP56), inference(cnf_transformation, [], [f147])).
fof(f147, plain, ((~ (e0 = op(e3, e4)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e3, e0))) | ~ sP56), inference(nnf_transformation, [], [f63])).
fof(f63, plain, ((~ (e0 = op(e3, e4)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e3, e0))) | ~ sP56), inference(usedef, [], [e63])).
fof(e63, plain, (sP56 <=> (~ (e0 = op(e3, e4)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP56])])).
fof(f1975, plain, (sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1974, f1021])).
fof(f1021, plain, ~ sP57, inference(subsumption_resolution, [], [f734, f678])).
fof(f734, plain, (~ (op(e3, e3) = unit) | ~ sP57), inference(definition_unfolding, [], [f377, f239])).
fof(f377, plain, (~ (e0 = op(e3, e3)) | ~ sP57), inference(cnf_transformation, [], [f146])).
fof(f146, plain, ((~ (e0 = op(e4, e3)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e0, e3))) | ~ sP57), inference(nnf_transformation, [], [f64])).
fof(f64, plain, ((~ (e0 = op(e4, e3)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e0, e3))) | ~ sP57), inference(usedef, [], [e64])).
fof(e64, plain, (sP57 <=> (~ (e0 = op(e4, e3)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP57])])).
fof(f1974, plain, (sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1973, f1019])).
fof(f1019, plain, ~ sP58, inference(subsumption_resolution, [], [f373, f233])).
fof(f373, plain, (~ (e1 = op(e3, e4)) | ~ sP58), inference(cnf_transformation, [], [f145])).
fof(f145, plain, ((~ (e1 = op(e3, e4)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e3, e0))) | ~ sP58), inference(nnf_transformation, [], [f65])).
fof(f65, plain, ((~ (e1 = op(e3, e4)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e3, e0))) | ~ sP58), inference(usedef, [], [e65])).
fof(e65, plain, (sP58 <=> (~ (e1 = op(e3, e4)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP58])])).
fof(f1973, plain, (sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1972, f1018])).
fof(f1018, plain, ~ sP59, inference(subsumption_resolution, [], [f366, f227])).
fof(f366, plain, (~ (e1 = op(e2, e3)) | ~ sP59), inference(cnf_transformation, [], [f144])).
fof(f144, plain, ((~ (e1 = op(e4, e3)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e0, e3))) | ~ sP59), inference(nnf_transformation, [], [f66])).
fof(f66, plain, ((~ (e1 = op(e4, e3)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e0, e3))) | ~ sP59), inference(usedef, [], [e66])).
fof(e66, plain, (sP59 <=> (~ (e1 = op(e4, e3)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP59])])).
fof(f1972, plain, (sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1971, f1015])).
fof(f1015, plain, ~ sP60, inference(subsumption_resolution, [], [f360, f230])).
fof(f360, plain, (~ (e2 = op(e3, e1)) | ~ sP60), inference(cnf_transformation, [], [f143])).
fof(f143, plain, ((~ (e2 = op(e3, e4)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e3, e0))) | ~ sP60), inference(nnf_transformation, [], [f67])).
fof(f67, plain, ((~ (e2 = op(e3, e4)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e3, e0))) | ~ sP60), inference(usedef, [], [e67])).
fof(e67, plain, (sP60 <=> (~ (e2 = op(e3, e4)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP60])])).
fof(f1971, plain, (sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1970, f1011])).
fof(f1011, plain, ~ sP61, inference(subsumption_resolution, [], [f358, f237])).
fof(f358, plain, (~ (e2 = op(e4, e3)) | ~ sP61), inference(cnf_transformation, [], [f142])).
fof(f142, plain, ((~ (e2 = op(e4, e3)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e0, e3))) | ~ sP61), inference(nnf_transformation, [], [f68])).
fof(f68, plain, ((~ (e2 = op(e4, e3)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e0, e3))) | ~ sP61), inference(usedef, [], [e68])).
fof(e68, plain, (sP61 <=> (~ (e2 = op(e4, e3)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP61])])).
fof(f1970, plain, (sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1969, f1010])).
fof(f1010, plain, ~ sP62, inference(subsumption_resolution, [], [f728, f679])).
fof(f728, plain, (~ (e3 = op(e3, unit)) | ~ sP62), inference(definition_unfolding, [], [f349, f239])).
fof(f349, plain, (~ (e3 = op(e3, e0)) | ~ sP62), inference(cnf_transformation, [], [f141])).
fof(f141, plain, ((~ (e3 = op(e3, e4)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e3, e0))) | ~ sP62), inference(nnf_transformation, [], [f69])).
fof(f69, plain, ((~ (e3 = op(e3, e4)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e3, e0))) | ~ sP62), inference(usedef, [], [e69])).
fof(e69, plain, (sP62 <=> (~ (e3 = op(e3, e4)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP62])])).
fof(f1969, plain, (sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1968, f1005])).
fof(f1005, plain, ~ sP63, inference(subsumption_resolution, [], [f727, f685])).
fof(f727, plain, (~ (e3 = op(unit, e3)) | ~ sP63), inference(definition_unfolding, [], [f344, f239])).
fof(f344, plain, (~ (e3 = op(e0, e3)) | ~ sP63), inference(cnf_transformation, [], [f140])).
fof(f140, plain, ((~ (e3 = op(e4, e3)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e0, e3))) | ~ sP63), inference(nnf_transformation, [], [f70])).
fof(f70, plain, ((~ (e3 = op(e4, e3)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e0, e3))) | ~ sP63), inference(usedef, [], [e70])).
fof(e70, plain, (sP63 <=> (~ (e3 = op(e4, e3)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP63])])).
fof(f1968, plain, (sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1967, f1000])).
fof(f1000, plain, ~ sP64, inference(subsumption_resolution, [], [f341, f231])).
fof(f341, plain, (~ (e4 = op(e3, e2)) | ~ sP64), inference(cnf_transformation, [], [f139])).
fof(f139, plain, ((~ (e4 = op(e3, e4)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e3, e0))) | ~ sP64), inference(nnf_transformation, [], [f71])).
fof(f71, plain, ((~ (e4 = op(e3, e4)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e3, e0))) | ~ sP64), inference(usedef, [], [e71])).
fof(e71, plain, (sP64 <=> (~ (e4 = op(e3, e4)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP64])])).
fof(f1967, plain, (sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1966, f997])).
fof(f997, plain, ~ sP65, inference(subsumption_resolution, [], [f335, f222])).
fof(f335, plain, (~ (e4 = op(e1, e3)) | ~ sP65), inference(cnf_transformation, [], [f138])).
fof(f138, plain, ((~ (e4 = op(e4, e3)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e0, e3))) | ~ sP65), inference(nnf_transformation, [], [f72])).
fof(f72, plain, ((~ (e4 = op(e4, e3)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e0, e3))) | ~ sP65), inference(usedef, [], [e72])).
fof(e72, plain, (sP65 <=> (~ (e4 = op(e4, e3)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP65])])).
fof(f1966, plain, (sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1965, f993])).
fof(f993, plain, ~ sP66, inference(subsumption_resolution, [], [f720, f676])).
fof(f720, plain, (~ (op(e4, e4) = unit) | ~ sP66), inference(definition_unfolding, [], [f333, f239])).
fof(f333, plain, (~ (e0 = op(e4, e4)) | ~ sP66), inference(cnf_transformation, [], [f137])).
fof(f137, plain, ((~ (e0 = op(e4, e4)) & ~ (e0 = op(e4, e3)) & ~ (e0 = op(e4, e2)) & ~ (e0 = op(e4, e1)) & ~ (e0 = op(e4, e0))) | ~ sP66), inference(nnf_transformation, [], [f73])).
fof(f73, plain, ((~ (e0 = op(e4, e4)) & ~ (e0 = op(e4, e3)) & ~ (e0 = op(e4, e2)) & ~ (e0 = op(e4, e1)) & ~ (e0 = op(e4, e0))) | ~ sP66), inference(usedef, [], [e73])).
fof(e73, plain, (sP66 <=> (~ (e0 = op(e4, e4)) & ~ (e0 = op(e4, e3)) & ~ (e0 = op(e4, e2)) & ~ (e0 = op(e4, e1)) & ~ (e0 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP66])])).
fof(f1965, plain, (sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1964, f992])).
fof(f992, plain, ~ sP67, inference(subsumption_resolution, [], [f715, f676])).
fof(f715, plain, (~ (op(e4, e4) = unit) | ~ sP67), inference(definition_unfolding, [], [f328, f239])).
fof(f328, plain, (~ (e0 = op(e4, e4)) | ~ sP67), inference(cnf_transformation, [], [f136])).
fof(f136, plain, ((~ (e0 = op(e4, e4)) & ~ (e0 = op(e3, e4)) & ~ (e0 = op(e2, e4)) & ~ (e0 = op(e1, e4)) & ~ (e0 = op(e0, e4))) | ~ sP67), inference(nnf_transformation, [], [f74])).
fof(f74, plain, ((~ (e0 = op(e4, e4)) & ~ (e0 = op(e3, e4)) & ~ (e0 = op(e2, e4)) & ~ (e0 = op(e1, e4)) & ~ (e0 = op(e0, e4))) | ~ sP67), inference(usedef, [], [e74])).
fof(e74, plain, (sP67 <=> (~ (e0 = op(e4, e4)) & ~ (e0 = op(e3, e4)) & ~ (e0 = op(e2, e4)) & ~ (e0 = op(e1, e4)) & ~ (e0 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP67])])).
fof(f1964, plain, (sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1963, f991])).
fof(f991, plain, ~ sP68, inference(subsumption_resolution, [], [f321, f236])).
fof(f321, plain, (~ (e1 = op(e4, e2)) | ~ sP68), inference(cnf_transformation, [], [f135])).
fof(f135, plain, ((~ (e1 = op(e4, e4)) & ~ (e1 = op(e4, e3)) & ~ (e1 = op(e4, e2)) & ~ (e1 = op(e4, e1)) & ~ (e1 = op(e4, e0))) | ~ sP68), inference(nnf_transformation, [], [f75])).
fof(f75, plain, ((~ (e1 = op(e4, e4)) & ~ (e1 = op(e4, e3)) & ~ (e1 = op(e4, e2)) & ~ (e1 = op(e4, e1)) & ~ (e1 = op(e4, e0))) | ~ sP68), inference(usedef, [], [e75])).
fof(e75, plain, (sP68 <=> (~ (e1 = op(e4, e4)) & ~ (e1 = op(e4, e3)) & ~ (e1 = op(e4, e2)) & ~ (e1 = op(e4, e1)) & ~ (e1 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP68])])).
fof(f1963, plain, (sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1962, f988])).
fof(f988, plain, ~ sP69, inference(subsumption_resolution, [], [f317, f233])).
fof(f317, plain, (~ (e1 = op(e3, e4)) | ~ sP69), inference(cnf_transformation, [], [f134])).
fof(f134, plain, ((~ (e1 = op(e4, e4)) & ~ (e1 = op(e3, e4)) & ~ (e1 = op(e2, e4)) & ~ (e1 = op(e1, e4)) & ~ (e1 = op(e0, e4))) | ~ sP69), inference(nnf_transformation, [], [f76])).
fof(f76, plain, ((~ (e1 = op(e4, e4)) & ~ (e1 = op(e3, e4)) & ~ (e1 = op(e2, e4)) & ~ (e1 = op(e1, e4)) & ~ (e1 = op(e0, e4))) | ~ sP69), inference(usedef, [], [e76])).
fof(e76, plain, (sP69 <=> (~ (e1 = op(e4, e4)) & ~ (e1 = op(e3, e4)) & ~ (e1 = op(e2, e4)) & ~ (e1 = op(e1, e4)) & ~ (e1 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP69])])).
fof(f1962, plain, (sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1961, f986])).
fof(f986, plain, ~ sP70, inference(subsumption_resolution, [], [f312, f237])).
fof(f312, plain, (~ (e2 = op(e4, e3)) | ~ sP70), inference(cnf_transformation, [], [f133])).
fof(f133, plain, ((~ (e2 = op(e4, e4)) & ~ (e2 = op(e4, e3)) & ~ (e2 = op(e4, e2)) & ~ (e2 = op(e4, e1)) & ~ (e2 = op(e4, e0))) | ~ sP70), inference(nnf_transformation, [], [f77])).
fof(f77, plain, ((~ (e2 = op(e4, e4)) & ~ (e2 = op(e4, e3)) & ~ (e2 = op(e4, e2)) & ~ (e2 = op(e4, e1)) & ~ (e2 = op(e4, e0))) | ~ sP70), inference(usedef, [], [e77])).
fof(e77, plain, (sP70 <=> (~ (e2 = op(e4, e4)) & ~ (e2 = op(e4, e3)) & ~ (e2 = op(e4, e2)) & ~ (e2 = op(e4, e1)) & ~ (e2 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP70])])).
fof(f1961, plain, (sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1960, f984])).
fof(f984, plain, ~ sP71, inference(subsumption_resolution, [], [f305, f223])).
fof(f305, plain, (~ (e2 = op(e1, e4)) | ~ sP71), inference(cnf_transformation, [], [f132])).
fof(f132, plain, ((~ (e2 = op(e4, e4)) & ~ (e2 = op(e3, e4)) & ~ (e2 = op(e2, e4)) & ~ (e2 = op(e1, e4)) & ~ (e2 = op(e0, e4))) | ~ sP71), inference(nnf_transformation, [], [f78])).
fof(f78, plain, ((~ (e2 = op(e4, e4)) & ~ (e2 = op(e3, e4)) & ~ (e2 = op(e2, e4)) & ~ (e2 = op(e1, e4)) & ~ (e2 = op(e0, e4))) | ~ sP71), inference(usedef, [], [e78])).
fof(e78, plain, (sP71 <=> (~ (e2 = op(e4, e4)) & ~ (e2 = op(e3, e4)) & ~ (e2 = op(e2, e4)) & ~ (e2 = op(e1, e4)) & ~ (e2 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP71])])).
fof(f1960, plain, (sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1959, f980])).
fof(f980, plain, ~ sP72, inference(subsumption_resolution, [], [f300, f235])).
fof(f300, plain, (~ (e3 = op(e4, e1)) | ~ sP72), inference(cnf_transformation, [], [f131])).
fof(f131, plain, ((~ (e3 = op(e4, e4)) & ~ (e3 = op(e4, e3)) & ~ (e3 = op(e4, e2)) & ~ (e3 = op(e4, e1)) & ~ (e3 = op(e4, e0))) | ~ sP72), inference(nnf_transformation, [], [f79])).
fof(f79, plain, ((~ (e3 = op(e4, e4)) & ~ (e3 = op(e4, e3)) & ~ (e3 = op(e4, e2)) & ~ (e3 = op(e4, e1)) & ~ (e3 = op(e4, e0))) | ~ sP72), inference(usedef, [], [e79])).
fof(e79, plain, (sP72 <=> (~ (e3 = op(e4, e4)) & ~ (e3 = op(e4, e3)) & ~ (e3 = op(e4, e2)) & ~ (e3 = op(e4, e1)) & ~ (e3 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP72])])).
fof(f1959, plain, (sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1958, f976])).
fof(f976, plain, ~ sP73, inference(subsumption_resolution, [], [f296, f228])).
fof(f296, plain, (~ (e3 = op(e2, e4)) | ~ sP73), inference(cnf_transformation, [], [f130])).
fof(f130, plain, ((~ (e3 = op(e4, e4)) & ~ (e3 = op(e3, e4)) & ~ (e3 = op(e2, e4)) & ~ (e3 = op(e1, e4)) & ~ (e3 = op(e0, e4))) | ~ sP73), inference(nnf_transformation, [], [f80])).
fof(f80, plain, ((~ (e3 = op(e4, e4)) & ~ (e3 = op(e3, e4)) & ~ (e3 = op(e2, e4)) & ~ (e3 = op(e1, e4)) & ~ (e3 = op(e0, e4))) | ~ sP73), inference(usedef, [], [e80])).
fof(e80, plain, (sP73 <=> (~ (e3 = op(e4, e4)) & ~ (e3 = op(e3, e4)) & ~ (e3 = op(e2, e4)) & ~ (e3 = op(e1, e4)) & ~ (e3 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP73])])).
fof(f1958, plain, (sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1957, f973])).
fof(f973, plain, ~ sP74, inference(subsumption_resolution, [], [f708, f677])).
fof(f708, plain, (~ (e4 = op(e4, unit)) | ~ sP74), inference(definition_unfolding, [], [f289, f239])).
fof(f289, plain, (~ (e4 = op(e4, e0)) | ~ sP74), inference(cnf_transformation, [], [f129])).
fof(f129, plain, ((~ (e4 = op(e4, e4)) & ~ (e4 = op(e4, e3)) & ~ (e4 = op(e4, e2)) & ~ (e4 = op(e4, e1)) & ~ (e4 = op(e4, e0))) | ~ sP74), inference(nnf_transformation, [], [f81])).
fof(f81, plain, ((~ (e4 = op(e4, e4)) & ~ (e4 = op(e4, e3)) & ~ (e4 = op(e4, e2)) & ~ (e4 = op(e4, e1)) & ~ (e4 = op(e4, e0))) | ~ sP74), inference(usedef, [], [e81])).
fof(e81, plain, (sP74 <=> (~ (e4 = op(e4, e4)) & ~ (e4 = op(e4, e3)) & ~ (e4 = op(e4, e2)) & ~ (e4 = op(e4, e1)) & ~ (e4 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP74])])).
fof(f1957, plain, (sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f899, f968])).
fof(f968, plain, ~ sP75, inference(subsumption_resolution, [], [f707, f684])).
fof(f707, plain, (~ (e4 = op(unit, e4)) | ~ sP75), inference(definition_unfolding, [], [f284, f239])).
fof(f284, plain, (~ (e4 = op(e0, e4)) | ~ sP75), inference(cnf_transformation, [], [f128])).
fof(f128, plain, ((~ (e4 = op(e4, e4)) & ~ (e4 = op(e3, e4)) & ~ (e4 = op(e2, e4)) & ~ (e4 = op(e1, e4)) & ~ (e4 = op(e0, e4))) | ~ sP75), inference(nnf_transformation, [], [f82])).
fof(f82, plain, ((~ (e4 = op(e4, e4)) & ~ (e4 = op(e3, e4)) & ~ (e4 = op(e2, e4)) & ~ (e4 = op(e1, e4)) & ~ (e4 = op(e0, e4))) | ~ sP75), inference(usedef, [], [e82])).
fof(e82, plain, (sP75 <=> (~ (e4 = op(e4, e4)) & ~ (e4 = op(e3, e4)) & ~ (e4 = op(e2, e4)) & ~ (e4 = op(e1, e4)) & ~ (e4 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP75])])).
fof(f899, plain, (sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(duplicate_literal_removal, [], [f898])).
fof(f898, plain, (sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(definition_unfolding, [], [f664, f239, f239, f239, f239])).
fof(f664, plain, (sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (e0 = op(e0, unit)) | ~ (e0 = op(unit, e0)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(cnf_transformation, [], [f105])).
fof(f105, plain, (sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (e0 = op(e0, unit)) | ~ (e0 = op(unit, e0)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | (~ (e4 = op(e4, e4)) & (e4 = op(e4, e4))) | (~ (e4 = op(e4, e3)) & (e3 = op(e4, e4))) | (~ (e4 = op(e4, e2)) & (e2 = op(e4, e4))) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(definition_folding, [], [f6, e104, e103, e102, e101, e100, e99, e98, e97, e96, e95, e94, e93, e92, e91, e90, e89, e88, e87, e86, e85, e84, e83, e82, e81, e80, e79, e78, e77, e76, e75, e74, e73, e72, e71, e70, e69, e68, e67, e66, e65, e64, e63, e62, e61, e60, e59, e58, e57, e56, e55, e54, e53, e52, e51, e50, e49, e48, e47, e46, e45, e44, e43, e42, e41, e40, e39, e38, e37, e36, e35, e34, e33, e32, e31, e30, e29, e28, e27, e26, e25, e24, e23, e22, e21, e20, e19, e18, e17, e16, e15, e14, e13, e12, e11, e10, e9, e8, e7])).
fof(f6, plain, ((~ (e4 = op(e4, e4)) & ~ (e4 = op(e3, e4)) & ~ (e4 = op(e2, e4)) & ~ (e4 = op(e1, e4)) & ~ (e4 = op(e0, e4))) | (~ (e4 = op(e4, e4)) & ~ (e4 = op(e4, e3)) & ~ (e4 = op(e4, e2)) & ~ (e4 = op(e4, e1)) & ~ (e4 = op(e4, e0))) | (~ (e3 = op(e4, e4)) & ~ (e3 = op(e3, e4)) & ~ (e3 = op(e2, e4)) & ~ (e3 = op(e1, e4)) & ~ (e3 = op(e0, e4))) | (~ (e3 = op(e4, e4)) & ~ (e3 = op(e4, e3)) & ~ (e3 = op(e4, e2)) & ~ (e3 = op(e4, e1)) & ~ (e3 = op(e4, e0))) | (~ (e2 = op(e4, e4)) & ~ (e2 = op(e3, e4)) & ~ (e2 = op(e2, e4)) & ~ (e2 = op(e1, e4)) & ~ (e2 = op(e0, e4))) | (~ (e2 = op(e4, e4)) & ~ (e2 = op(e4, e3)) & ~ (e2 = op(e4, e2)) & ~ (e2 = op(e4, e1)) & ~ (e2 = op(e4, e0))) | (~ (e1 = op(e4, e4)) & ~ (e1 = op(e3, e4)) & ~ (e1 = op(e2, e4)) & ~ (e1 = op(e1, e4)) & ~ (e1 = op(e0, e4))) | (~ (e1 = op(e4, e4)) & ~ (e1 = op(e4, e3)) & ~ (e1 = op(e4, e2)) & ~ (e1 = op(e4, e1)) & ~ (e1 = op(e4, e0))) | (~ (e0 = op(e4, e4)) & ~ (e0 = op(e3, e4)) & ~ (e0 = op(e2, e4)) & ~ (e0 = op(e1, e4)) & ~ (e0 = op(e0, e4))) | (~ (e0 = op(e4, e4)) & ~ (e0 = op(e4, e3)) & ~ (e0 = op(e4, e2)) & ~ (e0 = op(e4, e1)) & ~ (e0 = op(e4, e0))) | (~ (e4 = op(e4, e3)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e0, e3))) | (~ (e4 = op(e3, e4)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e3, e0))) | (~ (e3 = op(e4, e3)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e0, e3))) | (~ (e3 = op(e3, e4)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e3, e0))) | (~ (e2 = op(e4, e3)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e0, e3))) | (~ (e2 = op(e3, e4)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e3, e0))) | (~ (e1 = op(e4, e3)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e0, e3))) | (~ (e1 = op(e3, e4)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e3, e0))) | (~ (e0 = op(e4, e3)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e0, e3))) | (~ (e0 = op(e3, e4)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e3, e0))) | (~ (e4 = op(e4, e2)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e0, e2))) | (~ (e4 = op(e2, e4)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e2, e0))) | (~ (e3 = op(e4, e2)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e0, e2))) | (~ (e3 = op(e2, e4)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e2, e0))) | (~ (e2 = op(e4, e2)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e0, e2))) | (~ (e2 = op(e2, e4)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e2, e0))) | (~ (e1 = op(e4, e2)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e0, e2))) | (~ (e1 = op(e2, e4)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e2, e0))) | (~ (e0 = op(e4, e2)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e0, e2))) | (~ (e0 = op(e2, e4)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e2, e0))) | (~ (e4 = op(e4, e1)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e0, e1))) | (~ (e4 = op(e1, e4)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e1, e0))) | (~ (e3 = op(e4, e1)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e0, e1))) | (~ (e3 = op(e1, e4)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e1, e0))) | (~ (e2 = op(e4, e1)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e0, e1))) | (~ (e2 = op(e1, e4)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e1, e0))) | (~ (e1 = op(e4, e1)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e0, e1))) | (~ (e1 = op(e1, e4)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e1, e0))) | (~ (e0 = op(e4, e1)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e0, e1))) | (~ (e0 = op(e1, e4)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e1, e0))) | (~ (e4 = op(e4, e0)) & ~ (e4 = op(e3, e0)) & ~ (e4 = op(e2, e0)) & ~ (e4 = op(e1, e0)) & ~ (e4 = op(e0, e0))) | (~ (e4 = op(e0, e4)) & ~ (e4 = op(e0, e3)) & ~ (e4 = op(e0, e2)) & ~ (e4 = op(e0, e1)) & ~ (e4 = op(e0, e0))) | (~ (e3 = op(e4, e0)) & ~ (e3 = op(e3, e0)) & ~ (e3 = op(e2, e0)) & ~ (e3 = op(e1, e0)) & ~ (e3 = op(e0, e0))) | (~ (e3 = op(e0, e4)) & ~ (e3 = op(e0, e3)) & ~ (e3 = op(e0, e2)) & ~ (e3 = op(e0, e1)) & ~ (e3 = op(e0, e0))) | (~ (e2 = op(e4, e0)) & ~ (e2 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e2 = op(e1, e0)) & ~ (e2 = op(e0, e0))) | (~ (e2 = op(e0, e4)) & ~ (e2 = op(e0, e3)) & ~ (e2 = op(e0, e2)) & ~ (e2 = op(e0, e1)) & ~ (e2 = op(e0, e0))) | (~ (e1 = op(e4, e0)) & ~ (e1 = op(e3, e0)) & ~ (e1 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e1 = op(e0, e0))) | (~ (e1 = op(e0, e4)) & ~ (e1 = op(e0, e3)) & ~ (e1 = op(e0, e2)) & ~ (e1 = op(e0, e1)) & ~ (e1 = op(e0, e0))) | (~ (e0 = op(e4, e0)) & ~ (e0 = op(e3, e0)) & ~ (e0 = op(e2, e0)) & ~ (e0 = op(e1, e0)) & ~ (e0 = op(e0, e0))) | (~ (e0 = op(e0, e4)) & ~ (e0 = op(e0, e3)) & ~ (e0 = op(e0, e2)) & ~ (e0 = op(e0, e1)) & ~ (e0 = op(e0, e0))) | (~ (e4 = unit) & ~ (e3 = unit) & ~ (e2 = unit) & ~ (e1 = unit) & ~ (e0 = unit)) | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (e0 = op(e0, unit)) | ~ (e0 = op(unit, e0)) | (~ (e4 = op(e4, e4)) & ~ (e3 = op(e4, e4)) & ~ (e2 = op(e4, e4)) & ~ (e1 = op(e4, e4)) & ~ (e0 = op(e4, e4))) | (~ (e4 = op(e4, e3)) & ~ (e3 = op(e4, e3)) & ~ (e2 = op(e4, e3)) & ~ (e1 = op(e4, e3)) & ~ (e0 = op(e4, e3))) | (~ (e4 = op(e4, e2)) & ~ (e3 = op(e4, e2)) & ~ (e2 = op(e4, e2)) & ~ (e1 = op(e4, e2)) & ~ (e0 = op(e4, e2))) | (~ (e4 = op(e4, e1)) & ~ (e3 = op(e4, e1)) & ~ (e2 = op(e4, e1)) & ~ (e1 = op(e4, e1)) & ~ (e0 = op(e4, e1))) | (~ (e4 = op(e4, e0)) & ~ (e3 = op(e4, e0)) & ~ (e2 = op(e4, e0)) & ~ (e1 = op(e4, e0)) & ~ (e0 = op(e4, e0))) | (~ (e4 = op(e3, e4)) & ~ (e3 = op(e3, e4)) & ~ (e2 = op(e3, e4)) & ~ (e1 = op(e3, e4)) & ~ (e0 = op(e3, e4))) | (~ (e4 = op(e3, e3)) & ~ (e3 = op(e3, e3)) & ~ (e2 = op(e3, e3)) & ~ (e1 = op(e3, e3)) & ~ (e0 = op(e3, e3))) | (~ (e4 = op(e3, e2)) & ~ (e3 = op(e3, e2)) & ~ (e2 = op(e3, e2)) & ~ (e1 = op(e3, e2)) & ~ (e0 = op(e3, e2))) | (~ (e4 = op(e3, e1)) & ~ (e3 = op(e3, e1)) & ~ (e2 = op(e3, e1)) & ~ (e1 = op(e3, e1)) & ~ (e0 = op(e3, e1))) | (~ (e4 = op(e3, e0)) & ~ (e3 = op(e3, e0)) & ~ (e2 = op(e3, e0)) & ~ (e1 = op(e3, e0)) & ~ (e0 = op(e3, e0))) | (~ (e4 = op(e2, e4)) & ~ (e3 = op(e2, e4)) & ~ (e2 = op(e2, e4)) & ~ (e1 = op(e2, e4)) & ~ (e0 = op(e2, e4))) | (~ (e4 = op(e2, e3)) & ~ (e3 = op(e2, e3)) & ~ (e2 = op(e2, e3)) & ~ (e1 = op(e2, e3)) & ~ (e0 = op(e2, e3))) | (~ (e4 = op(e2, e2)) & ~ (e3 = op(e2, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e2, e2)) & ~ (e0 = op(e2, e2))) | (~ (e4 = op(e2, e1)) & ~ (e3 = op(e2, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e2, e1)) & ~ (e0 = op(e2, e1))) | (~ (e4 = op(e2, e0)) & ~ (e3 = op(e2, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e2, e0)) & ~ (e0 = op(e2, e0))) | (~ (e4 = op(e1, e4)) & ~ (e3 = op(e1, e4)) & ~ (e2 = op(e1, e4)) & ~ (e1 = op(e1, e4)) & ~ (e0 = op(e1, e4))) | (~ (e4 = op(e1, e3)) & ~ (e3 = op(e1, e3)) & ~ (e2 = op(e1, e3)) & ~ (e1 = op(e1, e3)) & ~ (e0 = op(e1, e3))) | (~ (e4 = op(e1, e2)) & ~ (e3 = op(e1, e2)) & ~ (e2 = op(e1, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e1, e2))) | (~ (e4 = op(e1, e1)) & ~ (e3 = op(e1, e1)) & ~ (e2 = op(e1, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e1, e1))) | (~ (e4 = op(e1, e0)) & ~ (e3 = op(e1, e0)) & ~ (e2 = op(e1, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e1, e0))) | (~ (e4 = op(e0, e4)) & ~ (e3 = op(e0, e4)) & ~ (e2 = op(e0, e4)) & ~ (e1 = op(e0, e4)) & ~ (e0 = op(e0, e4))) | (~ (e4 = op(e0, e3)) & ~ (e3 = op(e0, e3)) & ~ (e2 = op(e0, e3)) & ~ (e1 = op(e0, e3)) & ~ (e0 = op(e0, e3))) | (~ (e4 = op(e0, e2)) & ~ (e3 = op(e0, e2)) & ~ (e2 = op(e0, e2)) & ~ (e1 = op(e0, e2)) & ~ (e0 = op(e0, e2))) | (~ (e4 = op(e0, e1)) & ~ (e3 = op(e0, e1)) & ~ (e2 = op(e0, e1)) & ~ (e1 = op(e0, e1)) & ~ (e0 = op(e0, e1))) | (~ (e4 = op(e0, e0)) & ~ (e3 = op(e0, e0)) & ~ (e2 = op(e0, e0)) & ~ (e1 = op(e0, e0)) & ~ (e0 = op(e0, e0))) | (~ (e4 = op(e4, e4)) & (e4 = op(e4, e4))) | (~ (e4 = op(e4, e3)) & (e3 = op(e4, e4))) | (~ (e4 = op(e4, e2)) & (e2 = op(e4, e4))) | (~ (e4 = op(e4, e1)) & (e1 = op(e4, e4))) | (~ (e4 = op(e4, e0)) & (e0 = op(e4, e4))) | (~ (e3 = op(e3, e4)) & (e4 = op(e3, e3))) | (~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | (~ (e3 = op(e3, e2)) & (e2 = op(e3, e3))) | (~ (e3 = op(e3, e1)) & (e1 = op(e3, e3))) | (~ (e3 = op(e3, e0)) & (e0 = op(e3, e3))) | (~ (e2 = op(e2, e4)) & (e4 = op(e2, e2))) | (~ (e2 = op(e2, e3)) & (e3 = op(e2, e2))) | (~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | (~ (e2 = op(e2, e1)) & (e1 = op(e2, e2))) | (~ (e2 = op(e2, e0)) & (e0 = op(e2, e2))) | (~ (e1 = op(e1, e4)) & (e4 = op(e1, e1))) | (~ (e1 = op(e1, e3)) & (e3 = op(e1, e1))) | (~ (e1 = op(e1, e2)) & (e2 = op(e1, e1))) | (~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | (~ (e1 = op(e1, e0)) & (e0 = op(e1, e1))) | (~ (e0 = op(e0, e4)) & (e4 = op(e0, e0))) | (~ (e0 = op(e0, e3)) & (e3 = op(e0, e0))) | (~ (e0 = op(e0, e2)) & (e2 = op(e0, e0))) | (~ (e0 = op(e0, e1)) & (e1 = op(e0, e0))) | (~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)))), inference(ennf_transformation, [], [f5])).
fof(f5, plain, ~ (((e4 = op(e4, e4)) | (e4 = op(e3, e4)) | (e4 = op(e2, e4)) | (e4 = op(e1, e4)) | (e4 = op(e0, e4))) & ((e4 = op(e4, e4)) | (e4 = op(e4, e3)) | (e4 = op(e4, e2)) | (e4 = op(e4, e1)) | (e4 = op(e4, e0))) & ((e3 = op(e4, e4)) | (e3 = op(e3, e4)) | (e3 = op(e2, e4)) | (e3 = op(e1, e4)) | (e3 = op(e0, e4))) & ((e3 = op(e4, e4)) | (e3 = op(e4, e3)) | (e3 = op(e4, e2)) | (e3 = op(e4, e1)) | (e3 = op(e4, e0))) & ((e2 = op(e4, e4)) | (e2 = op(e3, e4)) | (e2 = op(e2, e4)) | (e2 = op(e1, e4)) | (e2 = op(e0, e4))) & ((e2 = op(e4, e4)) | (e2 = op(e4, e3)) | (e2 = op(e4, e2)) | (e2 = op(e4, e1)) | (e2 = op(e4, e0))) & ((e1 = op(e4, e4)) | (e1 = op(e3, e4)) | (e1 = op(e2, e4)) | (e1 = op(e1, e4)) | (e1 = op(e0, e4))) & ((e1 = op(e4, e4)) | (e1 = op(e4, e3)) | (e1 = op(e4, e2)) | (e1 = op(e4, e1)) | (e1 = op(e4, e0))) & ((e0 = op(e4, e4)) | (e0 = op(e3, e4)) | (e0 = op(e2, e4)) | (e0 = op(e1, e4)) | (e0 = op(e0, e4))) & ((e0 = op(e4, e4)) | (e0 = op(e4, e3)) | (e0 = op(e4, e2)) | (e0 = op(e4, e1)) | (e0 = op(e4, e0))) & ((e4 = op(e4, e3)) | (e4 = op(e3, e3)) | (e4 = op(e2, e3)) | (e4 = op(e1, e3)) | (e4 = op(e0, e3))) & ((e4 = op(e3, e4)) | (e4 = op(e3, e3)) | (e4 = op(e3, e2)) | (e4 = op(e3, e1)) | (e4 = op(e3, e0))) & ((e3 = op(e4, e3)) | (e3 = op(e3, e3)) | (e3 = op(e2, e3)) | (e3 = op(e1, e3)) | (e3 = op(e0, e3))) & ((e3 = op(e3, e4)) | (e3 = op(e3, e3)) | (e3 = op(e3, e2)) | (e3 = op(e3, e1)) | (e3 = op(e3, e0))) & ((e2 = op(e4, e3)) | (e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(e0, e3))) & ((e2 = op(e3, e4)) | (e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, e0))) & ((e1 = op(e4, e3)) | (e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(e0, e3))) & ((e1 = op(e3, e4)) | (e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, e0))) & ((e0 = op(e4, e3)) | (e0 = op(e3, e3)) | (e0 = op(e2, e3)) | (e0 = op(e1, e3)) | (e0 = op(e0, e3))) & ((e0 = op(e3, e4)) | (e0 = op(e3, e3)) | (e0 = op(e3, e2)) | (e0 = op(e3, e1)) | (e0 = op(e3, e0))) & ((e4 = op(e4, e2)) | (e4 = op(e3, e2)) | (e4 = op(e2, e2)) | (e4 = op(e1, e2)) | (e4 = op(e0, e2))) & ((e4 = op(e2, e4)) | (e4 = op(e2, e3)) | (e4 = op(e2, e2)) | (e4 = op(e2, e1)) | (e4 = op(e2, e0))) & ((e3 = op(e4, e2)) | (e3 = op(e3, e2)) | (e3 = op(e2, e2)) | (e3 = op(e1, e2)) | (e3 = op(e0, e2))) & ((e3 = op(e2, e4)) | (e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, e0))) & ((e2 = op(e4, e2)) | (e2 = op(e3, e2)) | (e2 = op(e2, e2)) | (e2 = op(e1, e2)) | (e2 = op(e0, e2))) & ((e2 = op(e2, e4)) | (e2 = op(e2, e3)) | (e2 = op(e2, e2)) | (e2 = op(e2, e1)) | (e2 = op(e2, e0))) & ((e1 = op(e4, e2)) | (e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(e0, e2))) & ((e1 = op(e2, e4)) | (e1 = op(e2, e3)) | (e1 = op(e2, e2)) | (e1 = op(e2, e1)) | (e1 = op(e2, e0))) & ((e0 = op(e4, e2)) | (e0 = op(e3, e2)) | (e0 = op(e2, e2)) | (e0 = op(e1, e2)) | (e0 = op(e0, e2))) & ((e0 = op(e2, e4)) | (e0 = op(e2, e3)) | (e0 = op(e2, e2)) | (e0 = op(e2, e1)) | (e0 = op(e2, e0))) & ((e4 = op(e4, e1)) | (e4 = op(e3, e1)) | (e4 = op(e2, e1)) | (e4 = op(e1, e1)) | (e4 = op(e0, e1))) & ((e4 = op(e1, e4)) | (e4 = op(e1, e3)) | (e4 = op(e1, e2)) | (e4 = op(e1, e1)) | (e4 = op(e1, e0))) & ((e3 = op(e4, e1)) | (e3 = op(e3, e1)) | (e3 = op(e2, e1)) | (e3 = op(e1, e1)) | (e3 = op(e0, e1))) & ((e3 = op(e1, e4)) | (e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, e0))) & ((e2 = op(e4, e1)) | (e2 = op(e3, e1)) | (e2 = op(e2, e1)) | (e2 = op(e1, e1)) | (e2 = op(e0, e1))) & ((e2 = op(e1, e4)) | (e2 = op(e1, e3)) | (e2 = op(e1, e2)) | (e2 = op(e1, e1)) | (e2 = op(e1, e0))) & ((e1 = op(e4, e1)) | (e1 = op(e3, e1)) | (e1 = op(e2, e1)) | (e1 = op(e1, e1)) | (e1 = op(e0, e1))) & ((e1 = op(e1, e4)) | (e1 = op(e1, e3)) | (e1 = op(e1, e2)) | (e1 = op(e1, e1)) | (e1 = op(e1, e0))) & ((e0 = op(e4, e1)) | (e0 = op(e3, e1)) | (e0 = op(e2, e1)) | (e0 = op(e1, e1)) | (e0 = op(e0, e1))) & ((e0 = op(e1, e4)) | (e0 = op(e1, e3)) | (e0 = op(e1, e2)) | (e0 = op(e1, e1)) | (e0 = op(e1, e0))) & ((e4 = op(e4, e0)) | (e4 = op(e3, e0)) | (e4 = op(e2, e0)) | (e4 = op(e1, e0)) | (e4 = op(e0, e0))) & ((e4 = op(e0, e4)) | (e4 = op(e0, e3)) | (e4 = op(e0, e2)) | (e4 = op(e0, e1)) | (e4 = op(e0, e0))) & ((e3 = op(e4, e0)) | (e3 = op(e3, e0)) | (e3 = op(e2, e0)) | (e3 = op(e1, e0)) | (e3 = op(e0, e0))) & ((e3 = op(e0, e4)) | (e3 = op(e0, e3)) | (e3 = op(e0, e2)) | (e3 = op(e0, e1)) | (e3 = op(e0, e0))) & ((e2 = op(e4, e0)) | (e2 = op(e3, e0)) | (e2 = op(e2, e0)) | (e2 = op(e1, e0)) | (e2 = op(e0, e0))) & ((e2 = op(e0, e4)) | (e2 = op(e0, e3)) | (e2 = op(e0, e2)) | (e2 = op(e0, e1)) | (e2 = op(e0, e0))) & ((e1 = op(e4, e0)) | (e1 = op(e3, e0)) | (e1 = op(e2, e0)) | (e1 = op(e1, e0)) | (e1 = op(e0, e0))) & ((e1 = op(e0, e4)) | (e1 = op(e0, e3)) | (e1 = op(e0, e2)) | (e1 = op(e0, e1)) | (e1 = op(e0, e0))) & ((e0 = op(e4, e0)) | (e0 = op(e3, e0)) | (e0 = op(e2, e0)) | (e0 = op(e1, e0)) | (e0 = op(e0, e0))) & ((e0 = op(e0, e4)) | (e0 = op(e0, e3)) | (e0 = op(e0, e2)) | (e0 = op(e0, e1)) | (e0 = op(e0, e0))) & ((e4 = unit) | (e3 = unit) | (e2 = unit) | (e1 = unit) | (e0 = unit)) & (e4 = op(e4, unit)) & (e4 = op(unit, e4)) & (e3 = op(e3, unit)) & (e3 = op(unit, e3)) & (e2 = op(e2, unit)) & (e2 = op(unit, e2)) & (e1 = op(e1, unit)) & (e1 = op(unit, e1)) & (e0 = op(e0, unit)) & (e0 = op(unit, e0)) & ((e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | (e1 = op(e4, e4)) | (e0 = op(e4, e4))) & ((e4 = op(e4, e3)) | (e3 = op(e4, e3)) | (e2 = op(e4, e3)) | (e1 = op(e4, e3)) | (e0 = op(e4, e3))) & ((e4 = op(e4, e2)) | (e3 = op(e4, e2)) | (e2 = op(e4, e2)) | (e1 = op(e4, e2)) | (e0 = op(e4, e2))) & ((e4 = op(e4, e1)) | (e3 = op(e4, e1)) | (e2 = op(e4, e1)) | (e1 = op(e4, e1)) | (e0 = op(e4, e1))) & ((e4 = op(e4, e0)) | (e3 = op(e4, e0)) | (e2 = op(e4, e0)) | (e1 = op(e4, e0)) | (e0 = op(e4, e0))) & ((e4 = op(e3, e4)) | (e3 = op(e3, e4)) | (e2 = op(e3, e4)) | (e1 = op(e3, e4)) | (e0 = op(e3, e4))) & ((e4 = op(e3, e3)) | (e3 = op(e3, e3)) | (e2 = op(e3, e3)) | (e1 = op(e3, e3)) | (e0 = op(e3, e3))) & ((e4 = op(e3, e2)) | (e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (e0 = op(e3, e2))) & ((e4 = op(e3, e1)) | (e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (e0 = op(e3, e1))) & ((e4 = op(e3, e0)) | (e3 = op(e3, e0)) | (e2 = op(e3, e0)) | (e1 = op(e3, e0)) | (e0 = op(e3, e0))) & ((e4 = op(e2, e4)) | (e3 = op(e2, e4)) | (e2 = op(e2, e4)) | (e1 = op(e2, e4)) | (e0 = op(e2, e4))) & ((e4 = op(e2, e3)) | (e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))) & ((e4 = op(e2, e2)) | (e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (e0 = op(e2, e2))) & ((e4 = op(e2, e1)) | (e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))) & ((e4 = op(e2, e0)) | (e3 = op(e2, e0)) | (e2 = op(e2, e0)) | (e1 = op(e2, e0)) | (e0 = op(e2, e0))) & ((e4 = op(e1, e4)) | (e3 = op(e1, e4)) | (e2 = op(e1, e4)) | (e1 = op(e1, e4)) | (e0 = op(e1, e4))) & ((e4 = op(e1, e3)) | (e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (e0 = op(e1, e3))) & ((e4 = op(e1, e2)) | (e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (e0 = op(e1, e2))) & ((e4 = op(e1, e1)) | (e3 = op(e1, e1)) | (e2 = op(e1, e1)) | (e1 = op(e1, e1)) | (e0 = op(e1, e1))) & ((e4 = op(e1, e0)) | (e3 = op(e1, e0)) | (e2 = op(e1, e0)) | (e1 = op(e1, e0)) | (e0 = op(e1, e0))) & ((e4 = op(e0, e4)) | (e3 = op(e0, e4)) | (e2 = op(e0, e4)) | (e1 = op(e0, e4)) | (e0 = op(e0, e4))) & ((e4 = op(e0, e3)) | (e3 = op(e0, e3)) | (e2 = op(e0, e3)) | (e1 = op(e0, e3)) | (e0 = op(e0, e3))) & ((e4 = op(e0, e2)) | (e3 = op(e0, e2)) | (e2 = op(e0, e2)) | (e1 = op(e0, e2)) | (e0 = op(e0, e2))) & ((e4 = op(e0, e1)) | (e3 = op(e0, e1)) | (e2 = op(e0, e1)) | (e1 = op(e0, e1)) | (e0 = op(e0, e1))) & ((e4 = op(e0, e0)) | (e3 = op(e0, e0)) | (e2 = op(e0, e0)) | (e1 = op(e0, e0)) | (e0 = op(e0, e0))) & ((e4 = op(e4, e4)) | ~ (e4 = op(e4, e4))) & ((e4 = op(e4, e3)) | ~ (e3 = op(e4, e4))) & ((e4 = op(e4, e2)) | ~ (e2 = op(e4, e4))) & ((e4 = op(e4, e1)) | ~ (e1 = op(e4, e4))) & ((e4 = op(e4, e0)) | ~ (e0 = op(e4, e4))) & ((e3 = op(e3, e4)) | ~ (e4 = op(e3, e3))) & ((e3 = op(e3, e3)) | ~ (e3 = op(e3, e3))) & ((e3 = op(e3, e2)) | ~ (e2 = op(e3, e3))) & ((e3 = op(e3, e1)) | ~ (e1 = op(e3, e3))) & ((e3 = op(e3, e0)) | ~ (e0 = op(e3, e3))) & ((e2 = op(e2, e4)) | ~ (e4 = op(e2, e2))) & ((e2 = op(e2, e3)) | ~ (e3 = op(e2, e2))) & ((e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & ((e2 = op(e2, e1)) | ~ (e1 = op(e2, e2))) & ((e2 = op(e2, e0)) | ~ (e0 = op(e2, e2))) & ((e1 = op(e1, e4)) | ~ (e4 = op(e1, e1))) & ((e1 = op(e1, e3)) | ~ (e3 = op(e1, e1))) & ((e1 = op(e1, e2)) | ~ (e2 = op(e1, e1))) & ((e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & ((e1 = op(e1, e0)) | ~ (e0 = op(e1, e1))) & ((e0 = op(e0, e4)) | ~ (e4 = op(e0, e0))) & ((e0 = op(e0, e3)) | ~ (e3 = op(e0, e0))) & ((e0 = op(e0, e2)) | ~ (e2 = op(e0, e0))) & ((e0 = op(e0, e1)) | ~ (e1 = op(e0, e0))) & ((e0 = op(e0, e0)) | ~ (e0 = op(e0, e0)))), inference(negated_conjecture, [], [f4])).
fof(f4, plain, ~ (((e4 = op(e4, e4)) | (e4 = op(e3, e4)) | (e4 = op(e2, e4)) | (e4 = op(e1, e4)) | (e4 = op(e0, e4))) & ((e4 = op(e4, e4)) | (e4 = op(e4, e3)) | (e4 = op(e4, e2)) | (e4 = op(e4, e1)) | (e4 = op(e4, e0))) & ((e3 = op(e4, e4)) | (e3 = op(e3, e4)) | (e3 = op(e2, e4)) | (e3 = op(e1, e4)) | (e3 = op(e0, e4))) & ((e3 = op(e4, e4)) | (e3 = op(e4, e3)) | (e3 = op(e4, e2)) | (e3 = op(e4, e1)) | (e3 = op(e4, e0))) & ((e2 = op(e4, e4)) | (e2 = op(e3, e4)) | (e2 = op(e2, e4)) | (e2 = op(e1, e4)) | (e2 = op(e0, e4))) & ((e2 = op(e4, e4)) | (e2 = op(e4, e3)) | (e2 = op(e4, e2)) | (e2 = op(e4, e1)) | (e2 = op(e4, e0))) & ((e1 = op(e4, e4)) | (e1 = op(e3, e4)) | (e1 = op(e2, e4)) | (e1 = op(e1, e4)) | (e1 = op(e0, e4))) & ((e1 = op(e4, e4)) | (e1 = op(e4, e3)) | (e1 = op(e4, e2)) | (e1 = op(e4, e1)) | (e1 = op(e4, e0))) & ((e0 = op(e4, e4)) | (e0 = op(e3, e4)) | (e0 = op(e2, e4)) | (e0 = op(e1, e4)) | (e0 = op(e0, e4))) & ((e0 = op(e4, e4)) | (e0 = op(e4, e3)) | (e0 = op(e4, e2)) | (e0 = op(e4, e1)) | (e0 = op(e4, e0))) & ((e4 = op(e4, e3)) | (e4 = op(e3, e3)) | (e4 = op(e2, e3)) | (e4 = op(e1, e3)) | (e4 = op(e0, e3))) & ((e4 = op(e3, e4)) | (e4 = op(e3, e3)) | (e4 = op(e3, e2)) | (e4 = op(e3, e1)) | (e4 = op(e3, e0))) & ((e3 = op(e4, e3)) | (e3 = op(e3, e3)) | (e3 = op(e2, e3)) | (e3 = op(e1, e3)) | (e3 = op(e0, e3))) & ((e3 = op(e3, e4)) | (e3 = op(e3, e3)) | (e3 = op(e3, e2)) | (e3 = op(e3, e1)) | (e3 = op(e3, e0))) & ((e2 = op(e4, e3)) | (e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(e0, e3))) & ((e2 = op(e3, e4)) | (e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, e0))) & ((e1 = op(e4, e3)) | (e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(e0, e3))) & ((e1 = op(e3, e4)) | (e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, e0))) & ((e0 = op(e4, e3)) | (e0 = op(e3, e3)) | (e0 = op(e2, e3)) | (e0 = op(e1, e3)) | (e0 = op(e0, e3))) & ((e0 = op(e3, e4)) | (e0 = op(e3, e3)) | (e0 = op(e3, e2)) | (e0 = op(e3, e1)) | (e0 = op(e3, e0))) & ((e4 = op(e4, e2)) | (e4 = op(e3, e2)) | (e4 = op(e2, e2)) | (e4 = op(e1, e2)) | (e4 = op(e0, e2))) & ((e4 = op(e2, e4)) | (e4 = op(e2, e3)) | (e4 = op(e2, e2)) | (e4 = op(e2, e1)) | (e4 = op(e2, e0))) & ((e3 = op(e4, e2)) | (e3 = op(e3, e2)) | (e3 = op(e2, e2)) | (e3 = op(e1, e2)) | (e3 = op(e0, e2))) & ((e3 = op(e2, e4)) | (e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, e0))) & ((e2 = op(e4, e2)) | (e2 = op(e3, e2)) | (e2 = op(e2, e2)) | (e2 = op(e1, e2)) | (e2 = op(e0, e2))) & ((e2 = op(e2, e4)) | (e2 = op(e2, e3)) | (e2 = op(e2, e2)) | (e2 = op(e2, e1)) | (e2 = op(e2, e0))) & ((e1 = op(e4, e2)) | (e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(e0, e2))) & ((e1 = op(e2, e4)) | (e1 = op(e2, e3)) | (e1 = op(e2, e2)) | (e1 = op(e2, e1)) | (e1 = op(e2, e0))) & ((e0 = op(e4, e2)) | (e0 = op(e3, e2)) | (e0 = op(e2, e2)) | (e0 = op(e1, e2)) | (e0 = op(e0, e2))) & ((e0 = op(e2, e4)) | (e0 = op(e2, e3)) | (e0 = op(e2, e2)) | (e0 = op(e2, e1)) | (e0 = op(e2, e0))) & ((e4 = op(e4, e1)) | (e4 = op(e3, e1)) | (e4 = op(e2, e1)) | (e4 = op(e1, e1)) | (e4 = op(e0, e1))) & ((e4 = op(e1, e4)) | (e4 = op(e1, e3)) | (e4 = op(e1, e2)) | (e4 = op(e1, e1)) | (e4 = op(e1, e0))) & ((e3 = op(e4, e1)) | (e3 = op(e3, e1)) | (e3 = op(e2, e1)) | (e3 = op(e1, e1)) | (e3 = op(e0, e1))) & ((e3 = op(e1, e4)) | (e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, e0))) & ((e2 = op(e4, e1)) | (e2 = op(e3, e1)) | (e2 = op(e2, e1)) | (e2 = op(e1, e1)) | (e2 = op(e0, e1))) & ((e2 = op(e1, e4)) | (e2 = op(e1, e3)) | (e2 = op(e1, e2)) | (e2 = op(e1, e1)) | (e2 = op(e1, e0))) & ((e1 = op(e4, e1)) | (e1 = op(e3, e1)) | (e1 = op(e2, e1)) | (e1 = op(e1, e1)) | (e1 = op(e0, e1))) & ((e1 = op(e1, e4)) | (e1 = op(e1, e3)) | (e1 = op(e1, e2)) | (e1 = op(e1, e1)) | (e1 = op(e1, e0))) & ((e0 = op(e4, e1)) | (e0 = op(e3, e1)) | (e0 = op(e2, e1)) | (e0 = op(e1, e1)) | (e0 = op(e0, e1))) & ((e0 = op(e1, e4)) | (e0 = op(e1, e3)) | (e0 = op(e1, e2)) | (e0 = op(e1, e1)) | (e0 = op(e1, e0))) & ((e4 = op(e4, e0)) | (e4 = op(e3, e0)) | (e4 = op(e2, e0)) | (e4 = op(e1, e0)) | (e4 = op(e0, e0))) & ((e4 = op(e0, e4)) | (e4 = op(e0, e3)) | (e4 = op(e0, e2)) | (e4 = op(e0, e1)) | (e4 = op(e0, e0))) & ((e3 = op(e4, e0)) | (e3 = op(e3, e0)) | (e3 = op(e2, e0)) | (e3 = op(e1, e0)) | (e3 = op(e0, e0))) & ((e3 = op(e0, e4)) | (e3 = op(e0, e3)) | (e3 = op(e0, e2)) | (e3 = op(e0, e1)) | (e3 = op(e0, e0))) & ((e2 = op(e4, e0)) | (e2 = op(e3, e0)) | (e2 = op(e2, e0)) | (e2 = op(e1, e0)) | (e2 = op(e0, e0))) & ((e2 = op(e0, e4)) | (e2 = op(e0, e3)) | (e2 = op(e0, e2)) | (e2 = op(e0, e1)) | (e2 = op(e0, e0))) & ((e1 = op(e4, e0)) | (e1 = op(e3, e0)) | (e1 = op(e2, e0)) | (e1 = op(e1, e0)) | (e1 = op(e0, e0))) & ((e1 = op(e0, e4)) | (e1 = op(e0, e3)) | (e1 = op(e0, e2)) | (e1 = op(e0, e1)) | (e1 = op(e0, e0))) & ((e0 = op(e4, e0)) | (e0 = op(e3, e0)) | (e0 = op(e2, e0)) | (e0 = op(e1, e0)) | (e0 = op(e0, e0))) & ((e0 = op(e0, e4)) | (e0 = op(e0, e3)) | (e0 = op(e0, e2)) | (e0 = op(e0, e1)) | (e0 = op(e0, e0))) & ((e4 = unit) | (e3 = unit) | (e2 = unit) | (e1 = unit) | (e0 = unit)) & (e4 = op(e4, unit)) & (e4 = op(unit, e4)) & (e3 = op(e3, unit)) & (e3 = op(unit, e3)) & (e2 = op(e2, unit)) & (e2 = op(unit, e2)) & (e1 = op(e1, unit)) & (e1 = op(unit, e1)) & (e0 = op(e0, unit)) & (e0 = op(unit, e0)) & ((e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | (e1 = op(e4, e4)) | (e0 = op(e4, e4))) & ((e4 = op(e4, e3)) | (e3 = op(e4, e3)) | (e2 = op(e4, e3)) | (e1 = op(e4, e3)) | (e0 = op(e4, e3))) & ((e4 = op(e4, e2)) | (e3 = op(e4, e2)) | (e2 = op(e4, e2)) | (e1 = op(e4, e2)) | (e0 = op(e4, e2))) & ((e4 = op(e4, e1)) | (e3 = op(e4, e1)) | (e2 = op(e4, e1)) | (e1 = op(e4, e1)) | (e0 = op(e4, e1))) & ((e4 = op(e4, e0)) | (e3 = op(e4, e0)) | (e2 = op(e4, e0)) | (e1 = op(e4, e0)) | (e0 = op(e4, e0))) & ((e4 = op(e3, e4)) | (e3 = op(e3, e4)) | (e2 = op(e3, e4)) | (e1 = op(e3, e4)) | (e0 = op(e3, e4))) & ((e4 = op(e3, e3)) | (e3 = op(e3, e3)) | (e2 = op(e3, e3)) | (e1 = op(e3, e3)) | (e0 = op(e3, e3))) & ((e4 = op(e3, e2)) | (e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (e0 = op(e3, e2))) & ((e4 = op(e3, e1)) | (e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (e0 = op(e3, e1))) & ((e4 = op(e3, e0)) | (e3 = op(e3, e0)) | (e2 = op(e3, e0)) | (e1 = op(e3, e0)) | (e0 = op(e3, e0))) & ((e4 = op(e2, e4)) | (e3 = op(e2, e4)) | (e2 = op(e2, e4)) | (e1 = op(e2, e4)) | (e0 = op(e2, e4))) & ((e4 = op(e2, e3)) | (e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))) & ((e4 = op(e2, e2)) | (e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (e0 = op(e2, e2))) & ((e4 = op(e2, e1)) | (e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))) & ((e4 = op(e2, e0)) | (e3 = op(e2, e0)) | (e2 = op(e2, e0)) | (e1 = op(e2, e0)) | (e0 = op(e2, e0))) & ((e4 = op(e1, e4)) | (e3 = op(e1, e4)) | (e2 = op(e1, e4)) | (e1 = op(e1, e4)) | (e0 = op(e1, e4))) & ((e4 = op(e1, e3)) | (e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (e0 = op(e1, e3))) & ((e4 = op(e1, e2)) | (e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (e0 = op(e1, e2))) & ((e4 = op(e1, e1)) | (e3 = op(e1, e1)) | (e2 = op(e1, e1)) | (e1 = op(e1, e1)) | (e0 = op(e1, e1))) & ((e4 = op(e1, e0)) | (e3 = op(e1, e0)) | (e2 = op(e1, e0)) | (e1 = op(e1, e0)) | (e0 = op(e1, e0))) & ((e4 = op(e0, e4)) | (e3 = op(e0, e4)) | (e2 = op(e0, e4)) | (e1 = op(e0, e4)) | (e0 = op(e0, e4))) & ((e4 = op(e0, e3)) | (e3 = op(e0, e3)) | (e2 = op(e0, e3)) | (e1 = op(e0, e3)) | (e0 = op(e0, e3))) & ((e4 = op(e0, e2)) | (e3 = op(e0, e2)) | (e2 = op(e0, e2)) | (e1 = op(e0, e2)) | (e0 = op(e0, e2))) & ((e4 = op(e0, e1)) | (e3 = op(e0, e1)) | (e2 = op(e0, e1)) | (e1 = op(e0, e1)) | (e0 = op(e0, e1))) & ((e4 = op(e0, e0)) | (e3 = op(e0, e0)) | (e2 = op(e0, e0)) | (e1 = op(e0, e0)) | (e0 = op(e0, e0))) & ((e4 = op(e4, e4)) | ~ (e4 = op(e4, e4))) & ((e4 = op(e4, e3)) | ~ (e3 = op(e4, e4))) & ((e4 = op(e4, e2)) | ~ (e2 = op(e4, e4))) & ((e4 = op(e4, e1)) | ~ (e1 = op(e4, e4))) & ((e4 = op(e4, e0)) | ~ (e0 = op(e4, e4))) & ((e3 = op(e3, e4)) | ~ (e4 = op(e3, e3))) & ((e3 = op(e3, e3)) | ~ (e3 = op(e3, e3))) & ((e3 = op(e3, e2)) | ~ (e2 = op(e3, e3))) & ((e3 = op(e3, e1)) | ~ (e1 = op(e3, e3))) & ((e3 = op(e3, e0)) | ~ (e0 = op(e3, e3))) & ((e2 = op(e2, e4)) | ~ (e4 = op(e2, e2))) & ((e2 = op(e2, e3)) | ~ (e3 = op(e2, e2))) & ((e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & ((e2 = op(e2, e1)) | ~ (e1 = op(e2, e2))) & ((e2 = op(e2, e0)) | ~ (e0 = op(e2, e2))) & ((e1 = op(e1, e4)) | ~ (e4 = op(e1, e1))) & ((e1 = op(e1, e3)) | ~ (e3 = op(e1, e1))) & ((e1 = op(e1, e2)) | ~ (e2 = op(e1, e1))) & ((e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & ((e1 = op(e1, e0)) | ~ (e0 = op(e1, e1))) & ((e0 = op(e0, e4)) | ~ (e4 = op(e0, e0))) & ((e0 = op(e0, e3)) | ~ (e3 = op(e0, e0))) & ((e0 = op(e0, e2)) | ~ (e2 = op(e0, e0))) & ((e0 = op(e0, e1)) | ~ (e1 = op(e0, e0))) & ((e0 = op(e0, e0)) | ~ (e0 = op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG091+1.p', co1)).
fof(f1632, plain, (spl98_1 | spl98_2 | ~ spl98_3), inference(avatar_split_clause, [], [f1619, f1629, f1625, f1621])).
fof(f1619, plain, (~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4))), inference(subsumption_resolution, [], [f1618, f963])).
fof(f1618, plain, (~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP76), inference(subsumption_resolution, [], [f1617, f962])).
fof(f1617, plain, (~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP77 | sP76), inference(subsumption_resolution, [], [f1616, f959])).
fof(f1616, plain, (~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1615, f956])).
fof(f1615, plain, (~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1614, f953])).
fof(f1614, plain, (~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1613, f950])).
fof(f1613, plain, (~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1612, f949])).
fof(f1612, plain, (~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1611, f946])).
fof(f1611, plain, (~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1610, f943])).
fof(f1610, plain, (~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1609, f940])).
fof(f1609, plain, (~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1608, f937])).
fof(f1608, plain, (~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1607, f936])).
fof(f1607, plain, (~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1606, f933])).
fof(f1606, plain, (~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1605, f930])).
fof(f1605, plain, (~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1604, f927])).
fof(f1604, plain, (~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1603, f924])).
fof(f1603, plain, (~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1602, f923])).
fof(f1602, plain, (~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1601, f920])).
fof(f1601, plain, (~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1600, f917])).
fof(f1600, plain, (~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1599, f914])).
fof(f1599, plain, (~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1598, f911])).
fof(f1598, plain, (~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1597, f910])).
fof(f1597, plain, (~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1596, f1188])).
fof(f1596, plain, (sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1595, f1183])).
fof(f1595, plain, (sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1594, f1179])).
fof(f1594, plain, (sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1593, f1176])).
fof(f1593, plain, (sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1592, f1174])).
fof(f1592, plain, (sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1591, f1173])).
fof(f1591, plain, (sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1590, f1169])).
fof(f1590, plain, (sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1589, f1164])).
fof(f1589, plain, (sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1588, f1162])).
fof(f1588, plain, (sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1587, f1161])).
fof(f1587, plain, (sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1586, f1158])).
fof(f1586, plain, (sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1585, f1155])).
fof(f1585, plain, (sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1584, f1154])).
fof(f1584, plain, (sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1583, f1149])).
fof(f1583, plain, (sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1582, f1145])).
fof(f1582, plain, (sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1581, f1143])).
fof(f1581, plain, (sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1580, f1141])).
fof(f1580, plain, (sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1579, f1138])).
fof(f1579, plain, (sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1578, f1137])).
fof(f1578, plain, (sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1577, f1132])).
fof(f1577, plain, (sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1576, f1128])).
fof(f1576, plain, (sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1575, f1127])).
fof(f1575, plain, (sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1574, f1125])).
fof(f1574, plain, (sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1573, f1121])).
fof(f1573, plain, (sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1572, f1118])).
fof(f1572, plain, (sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1571, f688])).
fof(f1571, plain, (~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1570, f687])).
fof(f1570, plain, (~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1569, f683])).
fof(f1569, plain, (~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1568, f686])).
fof(f1568, plain, (~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1567, f681])).
fof(f1567, plain, (~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1566, f685])).
fof(f1566, plain, (~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1565, f679])).
fof(f1565, plain, (~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1564, f684])).
fof(f1564, plain, (~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1563, f677])).
fof(f1563, plain, (~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1562, f907])).
fof(f1562, plain, (sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1561, f1113])).
fof(f1561, plain, (sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1560, f1108])).
fof(f1560, plain, (sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1559, f1103])).
fof(f1559, plain, (sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1558, f1099])).
fof(f1558, plain, (sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1557, f1095])).
fof(f1557, plain, (sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1556, f1092])).
fof(f1556, plain, (sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1555, f1089])).
fof(f1555, plain, (sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1554, f1087])).
fof(f1554, plain, (sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1553, f1085])).
fof(f1553, plain, (sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1552, f1084])).
fof(f1552, plain, (sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1551, f1083])).
fof(f1551, plain, (sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1550, f1079])).
fof(f1550, plain, (sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1549, f1075])).
fof(f1549, plain, (sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1548, f1070])).
fof(f1548, plain, (sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1547, f1065])).
fof(f1547, plain, (sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1546, f1064])).
fof(f1546, plain, (sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1545, f1062])).
fof(f1545, plain, (sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1544, f1059])).
fof(f1544, plain, (sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1543, f1058])).
fof(f1543, plain, (sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1542, f1056])).
fof(f1542, plain, (sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1541, f1053])).
fof(f1541, plain, (sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1540, f1050])).
fof(f1540, plain, (sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1539, f1047])).
fof(f1539, plain, (sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1538, f1045])).
fof(f1538, plain, (sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1537, f1044])).
fof(f1537, plain, (sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1536, f1039])).
fof(f1536, plain, (sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1535, f1034])).
fof(f1535, plain, (sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1534, f1033])).
fof(f1534, plain, (sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1533, f1029])).
fof(f1533, plain, (sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1532, f1025])).
fof(f1532, plain, (sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1531, f1023])).
fof(f1531, plain, (sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1530, f1021])).
fof(f1530, plain, (sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1529, f1019])).
fof(f1529, plain, (sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1528, f1018])).
fof(f1528, plain, (sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1527, f1015])).
fof(f1527, plain, (sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1526, f1011])).
fof(f1526, plain, (sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1525, f1010])).
fof(f1525, plain, (sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1524, f1005])).
fof(f1524, plain, (sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1523, f1000])).
fof(f1523, plain, (sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1522, f997])).
fof(f1522, plain, (sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1521, f993])).
fof(f1521, plain, (sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1520, f992])).
fof(f1520, plain, (sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1519, f991])).
fof(f1519, plain, (sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1518, f988])).
fof(f1518, plain, (sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1517, f986])).
fof(f1517, plain, (sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1516, f984])).
fof(f1516, plain, (sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1515, f980])).
fof(f1515, plain, (sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1514, f976])).
fof(f1514, plain, (sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f1513, f973])).
fof(f1513, plain, (sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(subsumption_resolution, [], [f903, f968])).
fof(f903, plain, (sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(duplicate_literal_removal, [], [f894])).
fof(f894, plain, (sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | ~ (unit = op(unit, unit)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(definition_unfolding, [], [f668, f239, f239, f239, f239])).
fof(f668, plain, (sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (e0 = op(e0, unit)) | ~ (e0 = op(unit, e0)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76), inference(cnf_transformation, [], [f105])).