fof(f3396, plain, $false, inference(avatar_sat_refutation, [], [f1397, f1400, f1448, f1450, f1452, f1466, f1481, f1665, f1789, f1943, f2062, f2100, f2106, f2171, f2214, f2268, f2309, f2323, f2328, f2346, f2348, f2350, f2351, f2352, f2368, f2417, f2450, f2452, f2458, f2517, f2632, f2742, f2935, f3035, f3037, f3080, f3166, f3257, f3334, f3357])).
fof(f3357, plain, (~ spl130_91 | ~ spl130_95), inference(avatar_contradiction_clause, [], [f3356])).
fof(f3356, plain, ($false | (~ spl130_91 | ~ spl130_95)), inference(subsumption_resolution, [], [f3353, f460])).
fof(f460, plain, ~ (e0 = e4), inference(cnf_transformation, [], [f5])).
fof(f5, plain, (~ (e3 = e4) & ~ (e2 = e4) & ~ (e2 = e3) & ~ (e1 = e4) & ~ (e1 = e3) & ~ (e1 = e2) & ~ (e0 = e4) & ~ (e0 = e3) & ~ (e0 = e2) & ~ (e0 = e1)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG062+1.p', ax5)).
fof(f3353, plain, ((e0 = e4) | (~ spl130_91 | ~ spl130_95)), inference(backward_demodulation, [], [f1249, f1233])).
fof(f1233, plain, ((e0 = op(e1, e1)) | ~ spl130_91), inference(avatar_component_clause, [], [f1231])).
fof(f1231, plain, (spl130_91 <=> (e0 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl130_91])])).
fof(f1249, plain, ((e4 = op(e1, e1)) | ~ spl130_95), inference(avatar_component_clause, [], [f1247])).
fof(f1247, plain, (spl130_95 <=> (e4 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl130_95])])).
fof(f3334, plain, (~ spl130_3 | ~ spl130_130), inference(avatar_contradiction_clause, [], [f3333])).
fof(f3333, plain, ($false | (~ spl130_3 | ~ spl130_130)), inference(subsumption_resolution, [], [f3332, f465])).
fof(f465, plain, ~ (e2 = e4), inference(cnf_transformation, [], [f5])).
fof(f3332, plain, ((e2 = e4) | (~ spl130_3 | ~ spl130_130)), inference(forward_demodulation, [], [f3312, f863])).
fof(f863, plain, ((e2 = op(e4, e4)) | ~ spl130_3), inference(avatar_component_clause, [], [f861])).
fof(f861, plain, (spl130_3 <=> (e2 = op(e4, e4))), introduced(avatar_definition, [new_symbols(naming, [spl130_3])])).
fof(f3312, plain, ((e4 = op(e4, e4)) | ~ spl130_130), inference(backward_demodulation, [], [f305, f1396])).
fof(f1396, plain, ((e4 = unit) | ~ spl130_130), inference(avatar_component_clause, [], [f1394])).
fof(f1394, plain, (spl130_130 <=> (e4 = unit)), introduced(avatar_definition, [new_symbols(naming, [spl130_130])])).
fof(f305, plain, (e4 = op(e4, unit)), inference(cnf_transformation, [], [f2])).
fof(f2, plain, (((e4 = unit) | (e3 = unit) | (e2 = unit) | (e1 = unit) | (e0 = unit)) & (e4 = op(e4, unit)) & (e4 = op(unit, e4)) & (e3 = op(e3, unit)) & (e3 = op(unit, e3)) & (e2 = op(e2, unit)) & (e2 = op(unit, e2)) & (e1 = op(e1, unit)) & (e1 = op(unit, e1)) & (e0 = op(e0, unit)) & (e0 = op(unit, e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG062+1.p', ax2)).
fof(f3257, plain, (~ spl130_1 | ~ spl130_3), inference(avatar_contradiction_clause, [], [f3256])).
fof(f3256, plain, ($false | (~ spl130_1 | ~ spl130_3)), inference(subsumption_resolution, [], [f3255, f458])).
fof(f458, plain, ~ (e0 = e2), inference(cnf_transformation, [], [f5])).
fof(f3255, plain, ((e0 = e2) | (~ spl130_1 | ~ spl130_3)), inference(backward_demodulation, [], [f863, f855])).
fof(f855, plain, ((e0 = op(e4, e4)) | ~ spl130_1), inference(avatar_component_clause, [], [f853])).
fof(f853, plain, (spl130_1 <=> (e0 = op(e4, e4))), introduced(avatar_definition, [new_symbols(naming, [spl130_1])])).
fof(f3166, plain, (spl130_106 | ~ spl130_129), inference(avatar_split_clause, [], [f3009, f1390, f1294])).
fof(f1294, plain, (spl130_106 <=> (e0 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl130_106])])).
fof(f1390, plain, (spl130_129 <=> (e3 = unit)), introduced(avatar_definition, [new_symbols(naming, [spl130_129])])).
fof(f3009, plain, ((e0 = op(e0, e3)) | ~ spl130_129), inference(backward_demodulation, [], [f297, f1392])).
fof(f1392, plain, ((e3 = unit) | ~ spl130_129), inference(avatar_component_clause, [], [f1390])).
fof(f297, plain, (e0 = op(e0, unit)), inference(cnf_transformation, [], [f2])).
fof(f3080, plain, (~ spl130_29 | ~ spl130_49), inference(avatar_split_clause, [], [f3076, f1054, f970])).
fof(f970, plain, (spl130_29 <=> (e3 = op(e3, e4))), introduced(avatar_definition, [new_symbols(naming, [spl130_29])])).
fof(f1054, plain, (spl130_49 <=> (e3 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl130_49])])).
fof(f3076, plain, (~ (e3 = op(e3, e4)) | ~ spl130_49), inference(backward_demodulation, [], [f440, f1056])).
fof(f1056, plain, ((e3 = op(e3, e0)) | ~ spl130_49), inference(avatar_component_clause, [], [f1054])).
fof(f440, plain, ~ (op(e3, e0) = op(e3, e4)), inference(cnf_transformation, [], [f4])).
fof(f4, plain, (~ (op(e4, e3) = op(e4, e4)) & ~ (op(e4, e2) = op(e4, e4)) & ~ (op(e4, e2) = op(e4, e3)) & ~ (op(e4, e1) = op(e4, e4)) & ~ (op(e4, e1) = op(e4, e3)) & ~ (op(e4, e1) = op(e4, e2)) & ~ (op(e4, e0) = op(e4, e4)) & ~ (op(e4, e0) = op(e4, e3)) & ~ (op(e4, e0) = op(e4, e2)) & ~ (op(e4, e0) = op(e4, e1)) & ~ (op(e3, e3) = op(e3, e4)) & ~ (op(e3, e2) = op(e3, e4)) & ~ (op(e3, e2) = op(e3, e3)) & ~ (op(e3, e1) = op(e3, e4)) & ~ (op(e3, e1) = op(e3, e3)) & ~ (op(e3, e1) = op(e3, e2)) & ~ (op(e3, e0) = op(e3, e4)) & ~ (op(e3, e0) = op(e3, e3)) & ~ (op(e3, e0) = op(e3, e2)) & ~ (op(e3, e0) = op(e3, e1)) & ~ (op(e2, e3) = op(e2, e4)) & ~ (op(e2, e2) = op(e2, e4)) & ~ (op(e2, e2) = op(e2, e3)) & ~ (op(e2, e1) = op(e2, e4)) & ~ (op(e2, e1) = op(e2, e3)) & ~ (op(e2, e1) = op(e2, e2)) & ~ (op(e2, e0) = op(e2, e4)) & ~ (op(e2, e0) = op(e2, e3)) & ~ (op(e2, e0) = op(e2, e2)) & ~ (op(e2, e0) = op(e2, e1)) & ~ (op(e1, e3) = op(e1, e4)) & ~ (op(e1, e2) = op(e1, e4)) & ~ (op(e1, e2) = op(e1, e3)) & ~ (op(e1, e1) = op(e1, e4)) & ~ (op(e1, e1) = op(e1, e3)) & ~ (op(e1, e1) = op(e1, e2)) & ~ (op(e1, e0) = op(e1, e4)) & ~ (op(e1, e0) = op(e1, e3)) & ~ (op(e1, e0) = op(e1, e2)) & ~ (op(e1, e0) = op(e1, e1)) & ~ (op(e0, e3) = op(e0, e4)) & ~ (op(e0, e2) = op(e0, e4)) & ~ (op(e0, e2) = op(e0, e3)) & ~ (op(e0, e1) = op(e0, e4)) & ~ (op(e0, e1) = op(e0, e3)) & ~ (op(e0, e1) = op(e0, e2)) & ~ (op(e0, e0) = op(e0, e4)) & ~ (op(e0, e0) = op(e0, e3)) & ~ (op(e0, e0) = op(e0, e2)) & ~ (op(e0, e0) = op(e0, e1)) & ~ (op(e3, e4) = op(e4, e4)) & ~ (op(e2, e4) = op(e4, e4)) & ~ (op(e2, e4) = op(e3, e4)) & ~ (op(e1, e4) = op(e4, e4)) & ~ (op(e1, e4) = op(e3, e4)) & ~ (op(e1, e4) = op(e2, e4)) & ~ (op(e0, e4) = op(e4, e4)) & ~ (op(e0, e4) = op(e3, e4)) & ~ (op(e0, e4) = op(e2, e4)) & ~ (op(e0, e4) = op(e1, e4)) & ~ (op(e3, e3) = op(e4, e3)) & ~ (op(e2, e3) = op(e4, e3)) & ~ (op(e2, e3) = op(e3, e3)) & ~ (op(e1, e3) = op(e4, e3)) & ~ (op(e1, e3) = op(e3, e3)) & ~ (op(e1, e3) = op(e2, e3)) & ~ (op(e0, e3) = op(e4, e3)) & ~ (op(e0, e3) = op(e3, e3)) & ~ (op(e0, e3) = op(e2, e3)) & ~ (op(e0, e3) = op(e1, e3)) & ~ (op(e3, e2) = op(e4, e2)) & ~ (op(e2, e2) = op(e4, e2)) & ~ (op(e2, e2) = op(e3, e2)) & ~ (op(e1, e2) = op(e4, e2)) & ~ (op(e1, e2) = op(e3, e2)) & ~ (op(e1, e2) = op(e2, e2)) & ~ (op(e0, e2) = op(e4, e2)) & ~ (op(e0, e2) = op(e3, e2)) & ~ (op(e0, e2) = op(e2, e2)) & ~ (op(e0, e2) = op(e1, e2)) & ~ (op(e3, e1) = op(e4, e1)) & ~ (op(e2, e1) = op(e4, e1)) & ~ (op(e2, e1) = op(e3, e1)) & ~ (op(e1, e1) = op(e4, e1)) & ~ (op(e1, e1) = op(e3, e1)) & ~ (op(e1, e1) = op(e2, e1)) & ~ (op(e0, e1) = op(e4, e1)) & ~ (op(e0, e1) = op(e3, e1)) & ~ (op(e0, e1) = op(e2, e1)) & ~ (op(e0, e1) = op(e1, e1)) & ~ (op(e3, e0) = op(e4, e0)) & ~ (op(e2, e0) = op(e4, e0)) & ~ (op(e2, e0) = op(e3, e0)) & ~ (op(e1, e0) = op(e4, e0)) & ~ (op(e1, e0) = op(e3, e0)) & ~ (op(e1, e0) = op(e2, e0)) & ~ (op(e0, e0) = op(e4, e0)) & ~ (op(e0, e0) = op(e3, e0)) & ~ (op(e0, e0) = op(e2, e0)) & ~ (op(e0, e0) = op(e1, e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG062+1.p', ax4)).
fof(f3037, plain, (spl130_10 | ~ spl130_129), inference(avatar_split_clause, [], [f3017, f1390, f890])).
fof(f890, plain, (spl130_10 <=> (e4 = op(e4, e3))), introduced(avatar_definition, [new_symbols(naming, [spl130_10])])).
fof(f3017, plain, ((e4 = op(e4, e3)) | ~ spl130_129), inference(backward_demodulation, [], [f305, f1392])).
fof(f3035, plain, (~ spl130_31 | ~ spl130_129), inference(avatar_contradiction_clause, [], [f3034])).
fof(f3034, plain, ($false | (~ spl130_31 | ~ spl130_129)), inference(subsumption_resolution, [], [f3033, f459])).
fof(f459, plain, ~ (e0 = e3), inference(cnf_transformation, [], [f5])).
fof(f3033, plain, ((e0 = e3) | (~ spl130_31 | ~ spl130_129)), inference(forward_demodulation, [], [f3015, f981])).
fof(f981, plain, ((e0 = op(e3, e3)) | ~ spl130_31), inference(avatar_component_clause, [], [f979])).
fof(f979, plain, (spl130_31 <=> (e0 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl130_31])])).
fof(f3015, plain, ((e3 = op(e3, e3)) | ~ spl130_129), inference(backward_demodulation, [], [f303, f1392])).
fof(f303, plain, (e3 = op(e3, unit)), inference(cnf_transformation, [], [f2])).
fof(f2935, plain, (~ spl130_86 | ~ spl130_128), inference(avatar_contradiction_clause, [], [f2934])).
fof(f2934, plain, ($false | (~ spl130_86 | ~ spl130_128)), inference(subsumption_resolution, [], [f2933, f457])).
fof(f457, plain, ~ (e0 = e1), inference(cnf_transformation, [], [f5])).
fof(f2933, plain, ((e0 = e1) | (~ spl130_86 | ~ spl130_128)), inference(forward_demodulation, [], [f2919, f1212])).
fof(f1212, plain, ((e0 = op(e1, e2)) | ~ spl130_86), inference(avatar_component_clause, [], [f1210])).
fof(f1210, plain, (spl130_86 <=> (e0 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl130_86])])).
fof(f2919, plain, ((e1 = op(e1, e2)) | ~ spl130_128), inference(backward_demodulation, [], [f299, f1388])).
fof(f1388, plain, ((e2 = unit) | ~ spl130_128), inference(avatar_component_clause, [], [f1386])).
fof(f1386, plain, (spl130_128 <=> (e2 = unit)), introduced(avatar_definition, [new_symbols(naming, [spl130_128])])).
fof(f299, plain, (e1 = op(e1, unit)), inference(cnf_transformation, [], [f2])).
fof(f2742, plain, (~ spl130_86 | ~ spl130_127), inference(avatar_contradiction_clause, [], [f2741])).
fof(f2741, plain, ($false | (~ spl130_86 | ~ spl130_127)), inference(subsumption_resolution, [], [f2740, f458])).
fof(f2740, plain, ((e0 = e2) | (~ spl130_86 | ~ spl130_127)), inference(forward_demodulation, [], [f2727, f1212])).
fof(f2727, plain, ((e2 = op(e1, e2)) | ~ spl130_127), inference(backward_demodulation, [], [f300, f1384])).
fof(f1384, plain, ((e1 = unit) | ~ spl130_127), inference(avatar_component_clause, [], [f1382])).
fof(f1382, plain, (spl130_127 <=> (e1 = unit)), introduced(avatar_definition, [new_symbols(naming, [spl130_127])])).
fof(f300, plain, (e2 = op(unit, e2)), inference(cnf_transformation, [], [f2])).
fof(f2632, plain, (~ spl130_3 | ~ spl130_4), inference(avatar_contradiction_clause, [], [f2631])).
fof(f2631, plain, ($false | (~ spl130_3 | ~ spl130_4)), inference(subsumption_resolution, [], [f2630, f464])).
fof(f464, plain, ~ (e2 = e3), inference(cnf_transformation, [], [f5])).
fof(f2630, plain, ((e2 = e3) | (~ spl130_3 | ~ spl130_4)), inference(backward_demodulation, [], [f867, f863])).
fof(f867, plain, ((e3 = op(e4, e4)) | ~ spl130_4), inference(avatar_component_clause, [], [f865])).
fof(f865, plain, (spl130_4 <=> (e3 = op(e4, e4))), introduced(avatar_definition, [new_symbols(naming, [spl130_4])])).
fof(f2517, plain, (~ spl130_54 | ~ spl130_64), inference(avatar_split_clause, [], [f2513, f1117, f1075])).
fof(f1075, plain, (spl130_54 <=> (e3 = op(e2, e4))), introduced(avatar_definition, [new_symbols(naming, [spl130_54])])).
fof(f1117, plain, (spl130_64 <=> (e3 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl130_64])])).
fof(f2513, plain, (~ (e3 = op(e2, e4)) | ~ spl130_64), inference(backward_demodulation, [], [f435, f1119])).
fof(f1119, plain, ((e3 = op(e2, e2)) | ~ spl130_64), inference(avatar_component_clause, [], [f1117])).
fof(f435, plain, ~ (op(e2, e2) = op(e2, e4)), inference(cnf_transformation, [], [f4])).
fof(f2458, plain, (~ spl130_61 | ~ spl130_86), inference(avatar_split_clause, [], [f2453, f1210, f1105])).
fof(f1105, plain, (spl130_61 <=> (e0 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl130_61])])).
fof(f2453, plain, (~ (e0 = op(e2, e2)) | ~ spl130_86), inference(backward_demodulation, [], [f381, f1212])).
fof(f381, plain, ~ (op(e1, e2) = op(e2, e2)), inference(cnf_transformation, [], [f4])).
fof(f2452, plain, (spl130_3 | ~ spl130_95), inference(avatar_split_clause, [], [f2445, f1247, f861])).
fof(f2445, plain, ((e2 = op(e4, e4)) | ~ spl130_95), inference(backward_demodulation, [], [f468, f1249])).
fof(f468, plain, (e2 = op(op(e1, e1), op(e1, e1))), inference(cnf_transformation, [], [f6])).
fof(f6, plain, ((e4 = op(e1, e1)) & (e3 = op(op(op(e1, e1), op(e1, e1)), op(op(e1, e1), op(e1, e1)))) & (e2 = op(op(e1, e1), op(e1, e1))) & (e0 = op(e1, op(op(e1, e1), op(e1, e1))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG062+1.p', ax6)).
fof(f2450, plain, (~ spl130_85 | ~ spl130_95), inference(avatar_split_clause, [], [f2443, f1247, f1205])).
fof(f1205, plain, (spl130_85 <=> (e4 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl130_85])])).
fof(f2443, plain, (~ (e4 = op(e1, e3)) | ~ spl130_95), inference(backward_demodulation, [], [f422, f1249])).
fof(f422, plain, ~ (op(e1, e1) = op(e1, e3)), inference(cnf_transformation, [], [f4])).
fof(f2417, plain, (~ spl130_104 | ~ spl130_105), inference(avatar_contradiction_clause, [], [f2416])).
fof(f2416, plain, ($false | (~ spl130_104 | ~ spl130_105)), inference(subsumption_resolution, [], [f2415, f466])).
fof(f466, plain, ~ (e3 = e4), inference(cnf_transformation, [], [f5])).
fof(f2415, plain, ((e3 = e4) | (~ spl130_104 | ~ spl130_105)), inference(backward_demodulation, [], [f1291, f1287])).
fof(f1287, plain, ((e3 = op(e0, e4)) | ~ spl130_104), inference(avatar_component_clause, [], [f1285])).
fof(f1285, plain, (spl130_104 <=> (e3 = op(e0, e4))), introduced(avatar_definition, [new_symbols(naming, [spl130_104])])).
fof(f1291, plain, ((e4 = op(e0, e4)) | ~ spl130_105), inference(avatar_component_clause, [], [f1289])).
fof(f1289, plain, (spl130_105 <=> (e4 = op(e0, e4))), introduced(avatar_definition, [new_symbols(naming, [spl130_105])])).
fof(f2368, plain, (~ spl130_106 | ~ spl130_121), inference(avatar_split_clause, [], [f2359, f1357, f1294])).
fof(f1357, plain, (spl130_121 <=> (e0 = op(e0, e0))), introduced(avatar_definition, [new_symbols(naming, [spl130_121])])).
fof(f2359, plain, (~ (e0 = op(e0, e3)) | ~ spl130_121), inference(backward_demodulation, [], [f409, f1359])).
fof(f1359, plain, ((e0 = op(e0, e0)) | ~ spl130_121), inference(avatar_component_clause, [], [f1357])).
fof(f409, plain, ~ (op(e0, e0) = op(e0, e3)), inference(cnf_transformation, [], [f4])).
fof(f2352, plain, (spl130_25 | ~ spl130_126), inference(avatar_split_clause, [], [f2338, f1378, f953])).
fof(f953, plain, (spl130_25 <=> (e4 = op(e4, e0))), introduced(avatar_definition, [new_symbols(naming, [spl130_25])])).
fof(f1378, plain, (spl130_126 <=> (e0 = unit)), introduced(avatar_definition, [new_symbols(naming, [spl130_126])])).
fof(f2338, plain, ((e4 = op(e4, e0)) | ~ spl130_126), inference(backward_demodulation, [], [f305, f1380])).
fof(f1380, plain, ((e0 = unit) | ~ spl130_126), inference(avatar_component_clause, [], [f1378])).
fof(f2351, plain, (spl130_105 | ~ spl130_126), inference(avatar_split_clause, [], [f2337, f1378, f1289])).
fof(f2337, plain, ((e4 = op(e0, e4)) | ~ spl130_126), inference(backward_demodulation, [], [f304, f1380])).
fof(f304, plain, (e4 = op(unit, e4)), inference(cnf_transformation, [], [f2])).
fof(f2350, plain, (spl130_49 | ~ spl130_126), inference(avatar_split_clause, [], [f2336, f1378, f1054])).
fof(f2336, plain, ((e3 = op(e3, e0)) | ~ spl130_126), inference(backward_demodulation, [], [f303, f1380])).
fof(f2348, plain, (spl130_73 | ~ spl130_126), inference(avatar_split_clause, [], [f2334, f1378, f1155])).
fof(f1155, plain, (spl130_73 <=> (e2 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl130_73])])).
fof(f2334, plain, ((e2 = op(e2, e0)) | ~ spl130_126), inference(backward_demodulation, [], [f301, f1380])).
fof(f301, plain, (e2 = op(e2, unit)), inference(cnf_transformation, [], [f2])).
fof(f2346, plain, (spl130_97 | ~ spl130_126), inference(avatar_split_clause, [], [f2332, f1378, f1256])).
fof(f1256, plain, (spl130_97 <=> (e1 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl130_97])])).
fof(f2332, plain, ((e1 = op(e1, e0)) | ~ spl130_126), inference(backward_demodulation, [], [f299, f1380])).
fof(f2328, plain, (spl130_121 | spl130_91 | spl130_61 | spl130_31 | spl130_1), inference(avatar_split_clause, [], [f846, f853, f979, f1105, f1231, f1357])).
fof(f846, plain, ((e0 = op(e4, e4)) | (e0 = op(e3, e3)) | (e0 = op(e2, e2)) | (e0 = op(e1, e1)) | (e0 = op(e0, e0))), inference(cnf_transformation, [], [f140])).
fof(f140, plain, ((sP129 | sP128 | sP127 | sP126 | sP125) & ((e4 = op(e4, e4)) | (e4 = op(e3, e3)) | (e4 = op(e2, e2)) | (e4 = op(e1, e1)) | (op(e0, e0) = e4)) & ((e3 = op(e4, e4)) | (e3 = op(e3, e3)) | (e3 = op(e2, e2)) | (e3 = op(e1, e1)) | (op(e0, e0) = e3)) & ((e2 = op(e4, e4)) | (e2 = op(e3, e3)) | (e2 = op(e2, e2)) | (e2 = op(e1, e1)) | (op(e0, e0) = e2)) & ((e1 = op(e4, e4)) | (e1 = op(e3, e3)) | (e1 = op(e2, e2)) | (e1 = op(e1, e1)) | (op(e0, e0) = e1)) & ((e0 = op(e4, e4)) | (e0 = op(e3, e3)) | (e0 = op(e2, e2)) | (e0 = op(e1, e1)) | (e0 = op(e0, e0)))), inference(definition_folding, [], [f9, e139, e138, e137, e136, e135, e134, e133, e132, e131, e130, e129, e128, e127, e126, e125, e124, e123, e122, e121, e120, e119, e118, e117, e116, e115, e114, e113, e112, e111, e110, e109, e108, e107, e106, e105, e104, e103, e102, e101, e100, e99, e98, e97, e96, e95, e94, e93, e92, e91, e90, e89, e88, e87, e86, e85, e84, e83, e82, e81, e80, e79, e78, e77, e76, e75, e74, e73, e72, e71, e70, e69, e68, e67, e66, e65, e64, e63, e62, e61, e60, e59, e58, e57, e56, e55, e54, e53, e52, e51, e50, e49, e48, e47, e46, e45, e44, e43, e42, e41, e40, e39, e38, e37, e36, e35, e34, e33, e32, e31, e30, e29, e28, e27, e26, e25, e24, e23, e22, e21, e20, e19, e18, e17, e16, e15, e14, e13, e12, e11, e10])).
fof(f10, plain, ((~ (e0 = unit) & (e0 = op(e0, e0))) | ~ (e0 = op(e0, e0)) | ~ sP0), inference(usedef, [], [e10])).
fof(e10, plain, (sP0 <=> ((~ (e0 = unit) & (e0 = op(e0, e0))) | ~ (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f11, plain, ((~ (e0 = unit) & (op(e0, e0) = e1)) | ~ (e0 = op(e0, e1)) | ~ sP1), inference(usedef, [], [e11])).
fof(e11, plain, (sP1 <=> ((~ (e0 = unit) & (op(e0, e0) = e1)) | ~ (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f12, plain, ((~ (e0 = unit) & (op(e0, e0) = e2)) | ~ (e0 = op(e0, e2)) | ~ sP2), inference(usedef, [], [e12])).
fof(e12, plain, (sP2 <=> ((~ (e0 = unit) & (op(e0, e0) = e2)) | ~ (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f13, plain, ((~ (e0 = unit) & (op(e0, e0) = e3)) | ~ (e0 = op(e0, e3)) | ~ sP3), inference(usedef, [], [e13])).
fof(e13, plain, (sP3 <=> ((~ (e0 = unit) & (op(e0, e0) = e3)) | ~ (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f14, plain, ((~ (e0 = unit) & (op(e0, e0) = e4)) | ~ (e0 = op(e0, e4)) | ~ sP4), inference(usedef, [], [e14])).
fof(e14, plain, (sP4 <=> ((~ (e0 = unit) & (op(e0, e0) = e4)) | ~ (e0 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f15, plain, ((~ (e0 = unit) & (e0 = op(e1, e0))) | ~ (e0 = op(e1, e0)) | ~ sP5), inference(usedef, [], [e15])).
fof(e15, plain, (sP5 <=> ((~ (e0 = unit) & (e0 = op(e1, e0))) | ~ (e0 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f16, plain, ((~ (e0 = unit) & (e1 = op(e1, e0))) | ~ (e0 = op(e1, e1)) | ~ sP6), inference(usedef, [], [e16])).
fof(e16, plain, (sP6 <=> ((~ (e0 = unit) & (e1 = op(e1, e0))) | ~ (e0 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f17, plain, ((~ (e0 = unit) & (e2 = op(e1, e0))) | ~ (e0 = op(e1, e2)) | ~ sP7), inference(usedef, [], [e17])).
fof(e17, plain, (sP7 <=> ((~ (e0 = unit) & (e2 = op(e1, e0))) | ~ (e0 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f18, plain, ((~ (e0 = unit) & (e3 = op(e1, e0))) | ~ (e0 = op(e1, e3)) | ~ sP8), inference(usedef, [], [e18])).
fof(e18, plain, (sP8 <=> ((~ (e0 = unit) & (e3 = op(e1, e0))) | ~ (e0 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f19, plain, ((~ (e0 = unit) & (e4 = op(e1, e0))) | ~ (e0 = op(e1, e4)) | ~ sP9), inference(usedef, [], [e19])).
fof(e19, plain, (sP9 <=> ((~ (e0 = unit) & (e4 = op(e1, e0))) | ~ (e0 = op(e1, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f20, plain, ((~ (e0 = unit) & (e0 = op(e2, e0))) | ~ (e0 = op(e2, e0)) | ~ sP10), inference(usedef, [], [e20])).
fof(e20, plain, (sP10 <=> ((~ (e0 = unit) & (e0 = op(e2, e0))) | ~ (e0 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f21, plain, ((~ (e0 = unit) & (e1 = op(e2, e0))) | ~ (e0 = op(e2, e1)) | ~ sP11), inference(usedef, [], [e21])).
fof(e21, plain, (sP11 <=> ((~ (e0 = unit) & (e1 = op(e2, e0))) | ~ (e0 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f22, plain, ((~ (e0 = unit) & (e2 = op(e2, e0))) | ~ (e0 = op(e2, e2)) | ~ sP12), inference(usedef, [], [e22])).
fof(e22, plain, (sP12 <=> ((~ (e0 = unit) & (e2 = op(e2, e0))) | ~ (e0 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f23, plain, ((~ (e0 = unit) & (e3 = op(e2, e0))) | ~ (e0 = op(e2, e3)) | ~ sP13), inference(usedef, [], [e23])).
fof(e23, plain, (sP13 <=> ((~ (e0 = unit) & (e3 = op(e2, e0))) | ~ (e0 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f24, plain, ((~ (e0 = unit) & (e4 = op(e2, e0))) | ~ (e0 = op(e2, e4)) | ~ sP14), inference(usedef, [], [e24])).
fof(e24, plain, (sP14 <=> ((~ (e0 = unit) & (e4 = op(e2, e0))) | ~ (e0 = op(e2, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f25, plain, ((~ (e0 = unit) & (e0 = op(e3, e0))) | ~ (e0 = op(e3, e0)) | ~ sP15), inference(usedef, [], [e25])).
fof(e25, plain, (sP15 <=> ((~ (e0 = unit) & (e0 = op(e3, e0))) | ~ (e0 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f26, plain, ((~ (e0 = unit) & (e1 = op(e3, e0))) | ~ (e0 = op(e3, e1)) | ~ sP16), inference(usedef, [], [e26])).
fof(e26, plain, (sP16 <=> ((~ (e0 = unit) & (e1 = op(e3, e0))) | ~ (e0 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f27, plain, ((~ (e0 = unit) & (e2 = op(e3, e0))) | ~ (e0 = op(e3, e2)) | ~ sP17), inference(usedef, [], [e27])).
fof(e27, plain, (sP17 <=> ((~ (e0 = unit) & (e2 = op(e3, e0))) | ~ (e0 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f28, plain, ((~ (e0 = unit) & (e3 = op(e3, e0))) | ~ (e0 = op(e3, e3)) | ~ sP18), inference(usedef, [], [e28])).
fof(e28, plain, (sP18 <=> ((~ (e0 = unit) & (e3 = op(e3, e0))) | ~ (e0 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f29, plain, ((~ (e0 = unit) & (e4 = op(e3, e0))) | ~ (e0 = op(e3, e4)) | ~ sP19), inference(usedef, [], [e29])).
fof(e29, plain, (sP19 <=> ((~ (e0 = unit) & (e4 = op(e3, e0))) | ~ (e0 = op(e3, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP19])])).
fof(f30, plain, ((~ (e0 = unit) & (e0 = op(e4, e0))) | ~ (e0 = op(e4, e0)) | ~ sP20), inference(usedef, [], [e30])).
fof(e30, plain, (sP20 <=> ((~ (e0 = unit) & (e0 = op(e4, e0))) | ~ (e0 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP20])])).
fof(f31, plain, ((~ (e0 = unit) & (e1 = op(e4, e0))) | ~ (e0 = op(e4, e1)) | ~ sP21), inference(usedef, [], [e31])).
fof(e31, plain, (sP21 <=> ((~ (e0 = unit) & (e1 = op(e4, e0))) | ~ (e0 = op(e4, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP21])])).
fof(f32, plain, ((~ (e0 = unit) & (e2 = op(e4, e0))) | ~ (e0 = op(e4, e2)) | ~ sP22), inference(usedef, [], [e32])).
fof(e32, plain, (sP22 <=> ((~ (e0 = unit) & (e2 = op(e4, e0))) | ~ (e0 = op(e4, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP22])])).
fof(f33, plain, ((~ (e0 = unit) & (e3 = op(e4, e0))) | ~ (e0 = op(e4, e3)) | ~ sP23), inference(usedef, [], [e33])).
fof(e33, plain, (sP23 <=> ((~ (e0 = unit) & (e3 = op(e4, e0))) | ~ (e0 = op(e4, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP23])])).
fof(f34, plain, ((~ (e0 = unit) & (e4 = op(e4, e0))) | ~ (e0 = op(e4, e4)) | ~ sP24), inference(usedef, [], [e34])).
fof(e34, plain, (sP24 <=> ((~ (e0 = unit) & (e4 = op(e4, e0))) | ~ (e0 = op(e4, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP24])])).
fof(f35, plain, ((~ (e1 = unit) & (e0 = op(e0, e1))) | ~ (op(e0, e0) = e1) | ~ sP25), inference(usedef, [], [e35])).
fof(e35, plain, (sP25 <=> ((~ (e1 = unit) & (e0 = op(e0, e1))) | ~ (op(e0, e0) = e1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP25])])).
fof(f36, plain, ((~ (e1 = unit) & (e1 = op(e0, e1))) | ~ (e1 = op(e0, e1)) | ~ sP26), inference(usedef, [], [e36])).
fof(e36, plain, (sP26 <=> ((~ (e1 = unit) & (e1 = op(e0, e1))) | ~ (e1 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP26])])).
fof(f37, plain, ((~ (e1 = unit) & (e2 = op(e0, e1))) | ~ (e1 = op(e0, e2)) | ~ sP27), inference(usedef, [], [e37])).
fof(e37, plain, (sP27 <=> ((~ (e1 = unit) & (e2 = op(e0, e1))) | ~ (e1 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP27])])).
fof(f38, plain, ((~ (e1 = unit) & (e3 = op(e0, e1))) | ~ (e1 = op(e0, e3)) | ~ sP28), inference(usedef, [], [e38])).
fof(e38, plain, (sP28 <=> ((~ (e1 = unit) & (e3 = op(e0, e1))) | ~ (e1 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP28])])).
fof(f39, plain, ((~ (e1 = unit) & (e4 = op(e0, e1))) | ~ (e1 = op(e0, e4)) | ~ sP29), inference(usedef, [], [e39])).
fof(e39, plain, (sP29 <=> ((~ (e1 = unit) & (e4 = op(e0, e1))) | ~ (e1 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP29])])).
fof(f40, plain, ((~ (e1 = unit) & (e0 = op(e1, e1))) | ~ (e1 = op(e1, e0)) | ~ sP30), inference(usedef, [], [e40])).
fof(e40, plain, (sP30 <=> ((~ (e1 = unit) & (e0 = op(e1, e1))) | ~ (e1 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP30])])).
fof(f41, plain, ((~ (e1 = unit) & (e1 = op(e1, e1))) | ~ (e1 = op(e1, e1)) | ~ sP31), inference(usedef, [], [e41])).
fof(e41, plain, (sP31 <=> ((~ (e1 = unit) & (e1 = op(e1, e1))) | ~ (e1 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP31])])).
fof(f42, plain, ((~ (e1 = unit) & (e2 = op(e1, e1))) | ~ (e1 = op(e1, e2)) | ~ sP32), inference(usedef, [], [e42])).
fof(e42, plain, (sP32 <=> ((~ (e1 = unit) & (e2 = op(e1, e1))) | ~ (e1 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP32])])).
fof(f43, plain, ((~ (e1 = unit) & (e3 = op(e1, e1))) | ~ (e1 = op(e1, e3)) | ~ sP33), inference(usedef, [], [e43])).
fof(e43, plain, (sP33 <=> ((~ (e1 = unit) & (e3 = op(e1, e1))) | ~ (e1 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP33])])).
fof(f44, plain, ((~ (e1 = unit) & (e4 = op(e1, e1))) | ~ (e1 = op(e1, e4)) | ~ sP34), inference(usedef, [], [e44])).
fof(e44, plain, (sP34 <=> ((~ (e1 = unit) & (e4 = op(e1, e1))) | ~ (e1 = op(e1, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP34])])).
fof(f45, plain, ((~ (e1 = unit) & (e0 = op(e2, e1))) | ~ (e1 = op(e2, e0)) | ~ sP35), inference(usedef, [], [e45])).
fof(e45, plain, (sP35 <=> ((~ (e1 = unit) & (e0 = op(e2, e1))) | ~ (e1 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP35])])).
fof(f46, plain, ((~ (e1 = unit) & (e1 = op(e2, e1))) | ~ (e1 = op(e2, e1)) | ~ sP36), inference(usedef, [], [e46])).
fof(e46, plain, (sP36 <=> ((~ (e1 = unit) & (e1 = op(e2, e1))) | ~ (e1 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP36])])).
fof(f47, plain, ((~ (e1 = unit) & (e2 = op(e2, e1))) | ~ (e1 = op(e2, e2)) | ~ sP37), inference(usedef, [], [e47])).
fof(e47, plain, (sP37 <=> ((~ (e1 = unit) & (e2 = op(e2, e1))) | ~ (e1 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP37])])).
fof(f48, plain, ((~ (e1 = unit) & (e3 = op(e2, e1))) | ~ (e1 = op(e2, e3)) | ~ sP38), inference(usedef, [], [e48])).
fof(e48, plain, (sP38 <=> ((~ (e1 = unit) & (e3 = op(e2, e1))) | ~ (e1 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP38])])).
fof(f49, plain, ((~ (e1 = unit) & (e4 = op(e2, e1))) | ~ (e1 = op(e2, e4)) | ~ sP39), inference(usedef, [], [e49])).
fof(e49, plain, (sP39 <=> ((~ (e1 = unit) & (e4 = op(e2, e1))) | ~ (e1 = op(e2, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP39])])).
fof(f50, plain, ((~ (e1 = unit) & (e0 = op(e3, e1))) | ~ (e1 = op(e3, e0)) | ~ sP40), inference(usedef, [], [e50])).
fof(e50, plain, (sP40 <=> ((~ (e1 = unit) & (e0 = op(e3, e1))) | ~ (e1 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP40])])).
fof(f51, plain, ((~ (e1 = unit) & (e1 = op(e3, e1))) | ~ (e1 = op(e3, e1)) | ~ sP41), inference(usedef, [], [e51])).
fof(e51, plain, (sP41 <=> ((~ (e1 = unit) & (e1 = op(e3, e1))) | ~ (e1 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP41])])).
fof(f52, plain, ((~ (e1 = unit) & (e2 = op(e3, e1))) | ~ (e1 = op(e3, e2)) | ~ sP42), inference(usedef, [], [e52])).
fof(e52, plain, (sP42 <=> ((~ (e1 = unit) & (e2 = op(e3, e1))) | ~ (e1 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP42])])).
fof(f53, plain, ((~ (e1 = unit) & (e3 = op(e3, e1))) | ~ (e1 = op(e3, e3)) | ~ sP43), inference(usedef, [], [e53])).
fof(e53, plain, (sP43 <=> ((~ (e1 = unit) & (e3 = op(e3, e1))) | ~ (e1 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP43])])).
fof(f54, plain, ((~ (e1 = unit) & (e4 = op(e3, e1))) | ~ (e1 = op(e3, e4)) | ~ sP44), inference(usedef, [], [e54])).
fof(e54, plain, (sP44 <=> ((~ (e1 = unit) & (e4 = op(e3, e1))) | ~ (e1 = op(e3, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP44])])).
fof(f55, plain, ((~ (e1 = unit) & (e0 = op(e4, e1))) | ~ (e1 = op(e4, e0)) | ~ sP45), inference(usedef, [], [e55])).
fof(e55, plain, (sP45 <=> ((~ (e1 = unit) & (e0 = op(e4, e1))) | ~ (e1 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP45])])).
fof(f56, plain, ((~ (e1 = unit) & (e1 = op(e4, e1))) | ~ (e1 = op(e4, e1)) | ~ sP46), inference(usedef, [], [e56])).
fof(e56, plain, (sP46 <=> ((~ (e1 = unit) & (e1 = op(e4, e1))) | ~ (e1 = op(e4, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP46])])).
fof(f57, plain, ((~ (e1 = unit) & (e2 = op(e4, e1))) | ~ (e1 = op(e4, e2)) | ~ sP47), inference(usedef, [], [e57])).
fof(e57, plain, (sP47 <=> ((~ (e1 = unit) & (e2 = op(e4, e1))) | ~ (e1 = op(e4, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP47])])).
fof(f58, plain, ((~ (e1 = unit) & (e3 = op(e4, e1))) | ~ (e1 = op(e4, e3)) | ~ sP48), inference(usedef, [], [e58])).
fof(e58, plain, (sP48 <=> ((~ (e1 = unit) & (e3 = op(e4, e1))) | ~ (e1 = op(e4, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP48])])).
fof(f59, plain, ((~ (e1 = unit) & (e4 = op(e4, e1))) | ~ (e1 = op(e4, e4)) | ~ sP49), inference(usedef, [], [e59])).
fof(e59, plain, (sP49 <=> ((~ (e1 = unit) & (e4 = op(e4, e1))) | ~ (e1 = op(e4, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP49])])).
fof(f60, plain, ((~ (e2 = unit) & (e0 = op(e0, e2))) | ~ (op(e0, e0) = e2) | ~ sP50), inference(usedef, [], [e60])).
fof(e60, plain, (sP50 <=> ((~ (e2 = unit) & (e0 = op(e0, e2))) | ~ (op(e0, e0) = e2))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP50])])).
fof(f61, plain, ((~ (e2 = unit) & (e1 = op(e0, e2))) | ~ (e2 = op(e0, e1)) | ~ sP51), inference(usedef, [], [e61])).
fof(e61, plain, (sP51 <=> ((~ (e2 = unit) & (e1 = op(e0, e2))) | ~ (e2 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP51])])).
fof(f62, plain, ((~ (e2 = unit) & (e2 = op(e0, e2))) | ~ (e2 = op(e0, e2)) | ~ sP52), inference(usedef, [], [e62])).
fof(e62, plain, (sP52 <=> ((~ (e2 = unit) & (e2 = op(e0, e2))) | ~ (e2 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP52])])).
fof(f63, plain, ((~ (e2 = unit) & (e3 = op(e0, e2))) | ~ (e2 = op(e0, e3)) | ~ sP53), inference(usedef, [], [e63])).
fof(e63, plain, (sP53 <=> ((~ (e2 = unit) & (e3 = op(e0, e2))) | ~ (e2 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP53])])).
fof(f64, plain, ((~ (e2 = unit) & (e4 = op(e0, e2))) | ~ (e2 = op(e0, e4)) | ~ sP54), inference(usedef, [], [e64])).
fof(e64, plain, (sP54 <=> ((~ (e2 = unit) & (e4 = op(e0, e2))) | ~ (e2 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP54])])).
fof(f65, plain, ((~ (e2 = unit) & (e0 = op(e1, e2))) | ~ (e2 = op(e1, e0)) | ~ sP55), inference(usedef, [], [e65])).
fof(e65, plain, (sP55 <=> ((~ (e2 = unit) & (e0 = op(e1, e2))) | ~ (e2 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP55])])).
fof(f66, plain, ((~ (e2 = unit) & (e1 = op(e1, e2))) | ~ (e2 = op(e1, e1)) | ~ sP56), inference(usedef, [], [e66])).
fof(e66, plain, (sP56 <=> ((~ (e2 = unit) & (e1 = op(e1, e2))) | ~ (e2 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP56])])).
fof(f67, plain, ((~ (e2 = unit) & (e2 = op(e1, e2))) | ~ (e2 = op(e1, e2)) | ~ sP57), inference(usedef, [], [e67])).
fof(e67, plain, (sP57 <=> ((~ (e2 = unit) & (e2 = op(e1, e2))) | ~ (e2 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP57])])).
fof(f68, plain, ((~ (e2 = unit) & (e3 = op(e1, e2))) | ~ (e2 = op(e1, e3)) | ~ sP58), inference(usedef, [], [e68])).
fof(e68, plain, (sP58 <=> ((~ (e2 = unit) & (e3 = op(e1, e2))) | ~ (e2 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP58])])).
fof(f69, plain, ((~ (e2 = unit) & (e4 = op(e1, e2))) | ~ (e2 = op(e1, e4)) | ~ sP59), inference(usedef, [], [e69])).
fof(e69, plain, (sP59 <=> ((~ (e2 = unit) & (e4 = op(e1, e2))) | ~ (e2 = op(e1, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP59])])).
fof(f70, plain, ((~ (e2 = unit) & (e0 = op(e2, e2))) | ~ (e2 = op(e2, e0)) | ~ sP60), inference(usedef, [], [e70])).
fof(e70, plain, (sP60 <=> ((~ (e2 = unit) & (e0 = op(e2, e2))) | ~ (e2 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP60])])).
fof(f71, plain, ((~ (e2 = unit) & (e1 = op(e2, e2))) | ~ (e2 = op(e2, e1)) | ~ sP61), inference(usedef, [], [e71])).
fof(e71, plain, (sP61 <=> ((~ (e2 = unit) & (e1 = op(e2, e2))) | ~ (e2 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP61])])).
fof(f72, plain, ((~ (e2 = unit) & (e2 = op(e2, e2))) | ~ (e2 = op(e2, e2)) | ~ sP62), inference(usedef, [], [e72])).
fof(e72, plain, (sP62 <=> ((~ (e2 = unit) & (e2 = op(e2, e2))) | ~ (e2 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP62])])).
fof(f73, plain, ((~ (e2 = unit) & (e3 = op(e2, e2))) | ~ (e2 = op(e2, e3)) | ~ sP63), inference(usedef, [], [e73])).
fof(e73, plain, (sP63 <=> ((~ (e2 = unit) & (e3 = op(e2, e2))) | ~ (e2 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP63])])).
fof(f74, plain, ((~ (e2 = unit) & (e4 = op(e2, e2))) | ~ (e2 = op(e2, e4)) | ~ sP64), inference(usedef, [], [e74])).
fof(e74, plain, (sP64 <=> ((~ (e2 = unit) & (e4 = op(e2, e2))) | ~ (e2 = op(e2, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP64])])).
fof(f75, plain, ((~ (e2 = unit) & (e0 = op(e3, e2))) | ~ (e2 = op(e3, e0)) | ~ sP65), inference(usedef, [], [e75])).
fof(e75, plain, (sP65 <=> ((~ (e2 = unit) & (e0 = op(e3, e2))) | ~ (e2 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP65])])).
fof(f76, plain, ((~ (e2 = unit) & (e1 = op(e3, e2))) | ~ (e2 = op(e3, e1)) | ~ sP66), inference(usedef, [], [e76])).
fof(e76, plain, (sP66 <=> ((~ (e2 = unit) & (e1 = op(e3, e2))) | ~ (e2 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP66])])).
fof(f77, plain, ((~ (e2 = unit) & (e2 = op(e3, e2))) | ~ (e2 = op(e3, e2)) | ~ sP67), inference(usedef, [], [e77])).
fof(e77, plain, (sP67 <=> ((~ (e2 = unit) & (e2 = op(e3, e2))) | ~ (e2 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP67])])).
fof(f78, plain, ((~ (e2 = unit) & (e3 = op(e3, e2))) | ~ (e2 = op(e3, e3)) | ~ sP68), inference(usedef, [], [e78])).
fof(e78, plain, (sP68 <=> ((~ (e2 = unit) & (e3 = op(e3, e2))) | ~ (e2 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP68])])).
fof(f79, plain, ((~ (e2 = unit) & (e4 = op(e3, e2))) | ~ (e2 = op(e3, e4)) | ~ sP69), inference(usedef, [], [e79])).
fof(e79, plain, (sP69 <=> ((~ (e2 = unit) & (e4 = op(e3, e2))) | ~ (e2 = op(e3, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP69])])).
fof(f80, plain, ((~ (e2 = unit) & (e0 = op(e4, e2))) | ~ (e2 = op(e4, e0)) | ~ sP70), inference(usedef, [], [e80])).
fof(e80, plain, (sP70 <=> ((~ (e2 = unit) & (e0 = op(e4, e2))) | ~ (e2 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP70])])).
fof(f81, plain, ((~ (e2 = unit) & (e1 = op(e4, e2))) | ~ (e2 = op(e4, e1)) | ~ sP71), inference(usedef, [], [e81])).
fof(e81, plain, (sP71 <=> ((~ (e2 = unit) & (e1 = op(e4, e2))) | ~ (e2 = op(e4, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP71])])).
fof(f82, plain, ((~ (e2 = unit) & (e2 = op(e4, e2))) | ~ (e2 = op(e4, e2)) | ~ sP72), inference(usedef, [], [e82])).
fof(e82, plain, (sP72 <=> ((~ (e2 = unit) & (e2 = op(e4, e2))) | ~ (e2 = op(e4, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP72])])).
fof(f83, plain, ((~ (e2 = unit) & (e3 = op(e4, e2))) | ~ (e2 = op(e4, e3)) | ~ sP73), inference(usedef, [], [e83])).
fof(e83, plain, (sP73 <=> ((~ (e2 = unit) & (e3 = op(e4, e2))) | ~ (e2 = op(e4, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP73])])).
fof(f84, plain, ((~ (e2 = unit) & (e4 = op(e4, e2))) | ~ (e2 = op(e4, e4)) | ~ sP74), inference(usedef, [], [e84])).
fof(e84, plain, (sP74 <=> ((~ (e2 = unit) & (e4 = op(e4, e2))) | ~ (e2 = op(e4, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP74])])).
fof(f85, plain, ((~ (e3 = unit) & (e0 = op(e0, e3))) | ~ (op(e0, e0) = e3) | ~ sP75), inference(usedef, [], [e85])).
fof(e85, plain, (sP75 <=> ((~ (e3 = unit) & (e0 = op(e0, e3))) | ~ (op(e0, e0) = e3))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP75])])).
fof(f86, plain, ((~ (e3 = unit) & (e1 = op(e0, e3))) | ~ (e3 = op(e0, e1)) | ~ sP76), inference(usedef, [], [e86])).
fof(e86, plain, (sP76 <=> ((~ (e3 = unit) & (e1 = op(e0, e3))) | ~ (e3 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP76])])).
fof(f87, plain, ((~ (e3 = unit) & (e2 = op(e0, e3))) | ~ (e3 = op(e0, e2)) | ~ sP77), inference(usedef, [], [e87])).
fof(e87, plain, (sP77 <=> ((~ (e3 = unit) & (e2 = op(e0, e3))) | ~ (e3 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP77])])).
fof(f88, plain, ((~ (e3 = unit) & (e3 = op(e0, e3))) | ~ (e3 = op(e0, e3)) | ~ sP78), inference(usedef, [], [e88])).
fof(e88, plain, (sP78 <=> ((~ (e3 = unit) & (e3 = op(e0, e3))) | ~ (e3 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP78])])).
fof(f89, plain, ((~ (e3 = unit) & (e4 = op(e0, e3))) | ~ (e3 = op(e0, e4)) | ~ sP79), inference(usedef, [], [e89])).
fof(e89, plain, (sP79 <=> ((~ (e3 = unit) & (e4 = op(e0, e3))) | ~ (e3 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP79])])).
fof(f90, plain, ((~ (e3 = unit) & (e0 = op(e1, e3))) | ~ (e3 = op(e1, e0)) | ~ sP80), inference(usedef, [], [e90])).
fof(e90, plain, (sP80 <=> ((~ (e3 = unit) & (e0 = op(e1, e3))) | ~ (e3 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP80])])).
fof(f91, plain, ((~ (e3 = unit) & (e1 = op(e1, e3))) | ~ (e3 = op(e1, e1)) | ~ sP81), inference(usedef, [], [e91])).
fof(e91, plain, (sP81 <=> ((~ (e3 = unit) & (e1 = op(e1, e3))) | ~ (e3 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP81])])).
fof(f92, plain, ((~ (e3 = unit) & (e2 = op(e1, e3))) | ~ (e3 = op(e1, e2)) | ~ sP82), inference(usedef, [], [e92])).
fof(e92, plain, (sP82 <=> ((~ (e3 = unit) & (e2 = op(e1, e3))) | ~ (e3 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP82])])).
fof(f93, plain, ((~ (e3 = unit) & (e3 = op(e1, e3))) | ~ (e3 = op(e1, e3)) | ~ sP83), inference(usedef, [], [e93])).
fof(e93, plain, (sP83 <=> ((~ (e3 = unit) & (e3 = op(e1, e3))) | ~ (e3 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP83])])).
fof(f94, plain, ((~ (e3 = unit) & (e4 = op(e1, e3))) | ~ (e3 = op(e1, e4)) | ~ sP84), inference(usedef, [], [e94])).
fof(e94, plain, (sP84 <=> ((~ (e3 = unit) & (e4 = op(e1, e3))) | ~ (e3 = op(e1, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP84])])).
fof(f95, plain, ((~ (e3 = unit) & (e0 = op(e2, e3))) | ~ (e3 = op(e2, e0)) | ~ sP85), inference(usedef, [], [e95])).
fof(e95, plain, (sP85 <=> ((~ (e3 = unit) & (e0 = op(e2, e3))) | ~ (e3 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP85])])).
fof(f96, plain, ((~ (e3 = unit) & (e1 = op(e2, e3))) | ~ (e3 = op(e2, e1)) | ~ sP86), inference(usedef, [], [e96])).
fof(e96, plain, (sP86 <=> ((~ (e3 = unit) & (e1 = op(e2, e3))) | ~ (e3 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP86])])).
fof(f97, plain, ((~ (e3 = unit) & (e2 = op(e2, e3))) | ~ (e3 = op(e2, e2)) | ~ sP87), inference(usedef, [], [e97])).
fof(e97, plain, (sP87 <=> ((~ (e3 = unit) & (e2 = op(e2, e3))) | ~ (e3 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP87])])).
fof(f98, plain, ((~ (e3 = unit) & (e3 = op(e2, e3))) | ~ (e3 = op(e2, e3)) | ~ sP88), inference(usedef, [], [e98])).
fof(e98, plain, (sP88 <=> ((~ (e3 = unit) & (e3 = op(e2, e3))) | ~ (e3 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP88])])).
fof(f99, plain, ((~ (e3 = unit) & (e4 = op(e2, e3))) | ~ (e3 = op(e2, e4)) | ~ sP89), inference(usedef, [], [e99])).
fof(e99, plain, (sP89 <=> ((~ (e3 = unit) & (e4 = op(e2, e3))) | ~ (e3 = op(e2, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP89])])).
fof(f100, plain, ((~ (e3 = unit) & (e0 = op(e3, e3))) | ~ (e3 = op(e3, e0)) | ~ sP90), inference(usedef, [], [e100])).
fof(e100, plain, (sP90 <=> ((~ (e3 = unit) & (e0 = op(e3, e3))) | ~ (e3 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP90])])).
fof(f101, plain, ((~ (e3 = unit) & (e1 = op(e3, e3))) | ~ (e3 = op(e3, e1)) | ~ sP91), inference(usedef, [], [e101])).
fof(e101, plain, (sP91 <=> ((~ (e3 = unit) & (e1 = op(e3, e3))) | ~ (e3 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP91])])).
fof(f102, plain, ((~ (e3 = unit) & (e2 = op(e3, e3))) | ~ (e3 = op(e3, e2)) | ~ sP92), inference(usedef, [], [e102])).
fof(e102, plain, (sP92 <=> ((~ (e3 = unit) & (e2 = op(e3, e3))) | ~ (e3 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP92])])).
fof(f103, plain, ((~ (e3 = unit) & (e3 = op(e3, e3))) | ~ (e3 = op(e3, e3)) | ~ sP93), inference(usedef, [], [e103])).
fof(e103, plain, (sP93 <=> ((~ (e3 = unit) & (e3 = op(e3, e3))) | ~ (e3 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP93])])).
fof(f104, plain, ((~ (e3 = unit) & (e4 = op(e3, e3))) | ~ (e3 = op(e3, e4)) | ~ sP94), inference(usedef, [], [e104])).
fof(e104, plain, (sP94 <=> ((~ (e3 = unit) & (e4 = op(e3, e3))) | ~ (e3 = op(e3, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP94])])).
fof(f105, plain, ((~ (e3 = unit) & (e0 = op(e4, e3))) | ~ (e3 = op(e4, e0)) | ~ sP95), inference(usedef, [], [e105])).
fof(e105, plain, (sP95 <=> ((~ (e3 = unit) & (e0 = op(e4, e3))) | ~ (e3 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP95])])).
fof(f106, plain, ((~ (e3 = unit) & (e1 = op(e4, e3))) | ~ (e3 = op(e4, e1)) | ~ sP96), inference(usedef, [], [e106])).
fof(e106, plain, (sP96 <=> ((~ (e3 = unit) & (e1 = op(e4, e3))) | ~ (e3 = op(e4, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP96])])).
fof(f107, plain, ((~ (e3 = unit) & (e2 = op(e4, e3))) | ~ (e3 = op(e4, e2)) | ~ sP97), inference(usedef, [], [e107])).
fof(e107, plain, (sP97 <=> ((~ (e3 = unit) & (e2 = op(e4, e3))) | ~ (e3 = op(e4, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP97])])).
fof(f108, plain, ((~ (e3 = unit) & (e3 = op(e4, e3))) | ~ (e3 = op(e4, e3)) | ~ sP98), inference(usedef, [], [e108])).
fof(e108, plain, (sP98 <=> ((~ (e3 = unit) & (e3 = op(e4, e3))) | ~ (e3 = op(e4, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP98])])).
fof(f109, plain, ((~ (e3 = unit) & (e4 = op(e4, e3))) | ~ (e3 = op(e4, e4)) | ~ sP99), inference(usedef, [], [e109])).
fof(e109, plain, (sP99 <=> ((~ (e3 = unit) & (e4 = op(e4, e3))) | ~ (e3 = op(e4, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP99])])).
fof(f110, plain, ((~ (e4 = unit) & (e0 = op(e0, e4))) | ~ (op(e0, e0) = e4) | ~ sP100), inference(usedef, [], [e110])).
fof(e110, plain, (sP100 <=> ((~ (e4 = unit) & (e0 = op(e0, e4))) | ~ (op(e0, e0) = e4))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP100])])).
fof(f111, plain, ((~ (e4 = unit) & (e1 = op(e0, e4))) | ~ (e4 = op(e0, e1)) | ~ sP101), inference(usedef, [], [e111])).
fof(e111, plain, (sP101 <=> ((~ (e4 = unit) & (e1 = op(e0, e4))) | ~ (e4 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP101])])).
fof(f112, plain, ((~ (e4 = unit) & (e2 = op(e0, e4))) | ~ (e4 = op(e0, e2)) | ~ sP102), inference(usedef, [], [e112])).
fof(e112, plain, (sP102 <=> ((~ (e4 = unit) & (e2 = op(e0, e4))) | ~ (e4 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP102])])).
fof(f113, plain, ((~ (e4 = unit) & (e3 = op(e0, e4))) | ~ (e4 = op(e0, e3)) | ~ sP103), inference(usedef, [], [e113])).
fof(e113, plain, (sP103 <=> ((~ (e4 = unit) & (e3 = op(e0, e4))) | ~ (e4 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP103])])).
fof(f114, plain, ((~ (e4 = unit) & (e4 = op(e0, e4))) | ~ (e4 = op(e0, e4)) | ~ sP104), inference(usedef, [], [e114])).
fof(e114, plain, (sP104 <=> ((~ (e4 = unit) & (e4 = op(e0, e4))) | ~ (e4 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP104])])).
fof(f115, plain, ((~ (e4 = unit) & (e0 = op(e1, e4))) | ~ (e4 = op(e1, e0)) | ~ sP105), inference(usedef, [], [e115])).
fof(e115, plain, (sP105 <=> ((~ (e4 = unit) & (e0 = op(e1, e4))) | ~ (e4 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP105])])).
fof(f116, plain, ((~ (e4 = unit) & (e1 = op(e1, e4))) | ~ (e4 = op(e1, e1)) | ~ sP106), inference(usedef, [], [e116])).
fof(e116, plain, (sP106 <=> ((~ (e4 = unit) & (e1 = op(e1, e4))) | ~ (e4 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP106])])).
fof(f117, plain, ((~ (e4 = unit) & (e2 = op(e1, e4))) | ~ (e4 = op(e1, e2)) | ~ sP107), inference(usedef, [], [e117])).
fof(e117, plain, (sP107 <=> ((~ (e4 = unit) & (e2 = op(e1, e4))) | ~ (e4 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP107])])).
fof(f118, plain, ((~ (e4 = unit) & (e3 = op(e1, e4))) | ~ (e4 = op(e1, e3)) | ~ sP108), inference(usedef, [], [e118])).
fof(e118, plain, (sP108 <=> ((~ (e4 = unit) & (e3 = op(e1, e4))) | ~ (e4 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP108])])).
fof(f119, plain, ((~ (e4 = unit) & (e4 = op(e1, e4))) | ~ (e4 = op(e1, e4)) | ~ sP109), inference(usedef, [], [e119])).
fof(e119, plain, (sP109 <=> ((~ (e4 = unit) & (e4 = op(e1, e4))) | ~ (e4 = op(e1, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP109])])).
fof(f120, plain, ((~ (e4 = unit) & (e0 = op(e2, e4))) | ~ (e4 = op(e2, e0)) | ~ sP110), inference(usedef, [], [e120])).
fof(e120, plain, (sP110 <=> ((~ (e4 = unit) & (e0 = op(e2, e4))) | ~ (e4 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP110])])).
fof(f121, plain, ((~ (e4 = unit) & (e1 = op(e2, e4))) | ~ (e4 = op(e2, e1)) | ~ sP111), inference(usedef, [], [e121])).
fof(e121, plain, (sP111 <=> ((~ (e4 = unit) & (e1 = op(e2, e4))) | ~ (e4 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP111])])).
fof(f122, plain, ((~ (e4 = unit) & (e2 = op(e2, e4))) | ~ (e4 = op(e2, e2)) | ~ sP112), inference(usedef, [], [e122])).
fof(e122, plain, (sP112 <=> ((~ (e4 = unit) & (e2 = op(e2, e4))) | ~ (e4 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP112])])).
fof(f123, plain, ((~ (e4 = unit) & (e3 = op(e2, e4))) | ~ (e4 = op(e2, e3)) | ~ sP113), inference(usedef, [], [e123])).
fof(e123, plain, (sP113 <=> ((~ (e4 = unit) & (e3 = op(e2, e4))) | ~ (e4 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP113])])).
fof(f124, plain, ((~ (e4 = unit) & (e4 = op(e2, e4))) | ~ (e4 = op(e2, e4)) | ~ sP114), inference(usedef, [], [e124])).
fof(e124, plain, (sP114 <=> ((~ (e4 = unit) & (e4 = op(e2, e4))) | ~ (e4 = op(e2, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP114])])).
fof(f125, plain, ((~ (e4 = unit) & (e0 = op(e3, e4))) | ~ (e4 = op(e3, e0)) | ~ sP115), inference(usedef, [], [e125])).
fof(e125, plain, (sP115 <=> ((~ (e4 = unit) & (e0 = op(e3, e4))) | ~ (e4 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP115])])).
fof(f126, plain, ((~ (e4 = unit) & (e1 = op(e3, e4))) | ~ (e4 = op(e3, e1)) | ~ sP116), inference(usedef, [], [e126])).
fof(e126, plain, (sP116 <=> ((~ (e4 = unit) & (e1 = op(e3, e4))) | ~ (e4 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP116])])).
fof(f127, plain, ((~ (e4 = unit) & (e2 = op(e3, e4))) | ~ (e4 = op(e3, e2)) | ~ sP117), inference(usedef, [], [e127])).
fof(e127, plain, (sP117 <=> ((~ (e4 = unit) & (e2 = op(e3, e4))) | ~ (e4 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP117])])).
fof(f128, plain, ((~ (e4 = unit) & (e3 = op(e3, e4))) | ~ (e4 = op(e3, e3)) | ~ sP118), inference(usedef, [], [e128])).
fof(e128, plain, (sP118 <=> ((~ (e4 = unit) & (e3 = op(e3, e4))) | ~ (e4 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP118])])).
fof(f129, plain, ((~ (e4 = unit) & (e4 = op(e3, e4))) | ~ (e4 = op(e3, e4)) | ~ sP119), inference(usedef, [], [e129])).
fof(e129, plain, (sP119 <=> ((~ (e4 = unit) & (e4 = op(e3, e4))) | ~ (e4 = op(e3, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP119])])).
fof(f130, plain, ((~ (e4 = unit) & (e0 = op(e4, e4))) | ~ (e4 = op(e4, e0)) | ~ sP120), inference(usedef, [], [e130])).
fof(e130, plain, (sP120 <=> ((~ (e4 = unit) & (e0 = op(e4, e4))) | ~ (e4 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP120])])).
fof(f131, plain, ((~ (e4 = unit) & (e1 = op(e4, e4))) | ~ (e4 = op(e4, e1)) | ~ sP121), inference(usedef, [], [e131])).
fof(e131, plain, (sP121 <=> ((~ (e4 = unit) & (e1 = op(e4, e4))) | ~ (e4 = op(e4, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP121])])).
fof(f132, plain, ((~ (e4 = unit) & (e2 = op(e4, e4))) | ~ (e4 = op(e4, e2)) | ~ sP122), inference(usedef, [], [e132])).
fof(e132, plain, (sP122 <=> ((~ (e4 = unit) & (e2 = op(e4, e4))) | ~ (e4 = op(e4, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP122])])).
fof(f133, plain, ((~ (e4 = unit) & (e3 = op(e4, e4))) | ~ (e4 = op(e4, e3)) | ~ sP123), inference(usedef, [], [e133])).
fof(e133, plain, (sP123 <=> ((~ (e4 = unit) & (e3 = op(e4, e4))) | ~ (e4 = op(e4, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP123])])).
fof(f134, plain, ((~ (e4 = unit) & (e4 = op(e4, e4))) | ~ (e4 = op(e4, e4)) | ~ sP124), inference(usedef, [], [e134])).
fof(e134, plain, (sP124 <=> ((~ (e4 = unit) & (e4 = op(e4, e4))) | ~ (e4 = op(e4, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP124])])).
fof(f135, plain, ((sP24 & sP23 & sP22 & sP21 & sP20 & sP19 & sP18 & sP17 & sP16 & sP15 & sP14 & sP13 & sP12 & sP11 & sP10 & sP9 & sP8 & sP7 & sP6 & sP5 & sP4 & sP3 & sP2 & sP1 & sP0) | ~ sP125), inference(usedef, [], [e135])).
fof(e135, plain, (sP125 <=> (sP24 & sP23 & sP22 & sP21 & sP20 & sP19 & sP18 & sP17 & sP16 & sP15 & sP14 & sP13 & sP12 & sP11 & sP10 & sP9 & sP8 & sP7 & sP6 & sP5 & sP4 & sP3 & sP2 & sP1 & sP0)), introduced(predicate_definition_introduction, [new_symbols(naming, [sP125])])).
fof(f136, plain, ((sP49 & sP48 & sP47 & sP46 & sP45 & sP44 & sP43 & sP42 & sP41 & sP40 & sP39 & sP38 & sP37 & sP36 & sP35 & sP34 & sP33 & sP32 & sP31 & sP30 & sP29 & sP28 & sP27 & sP26 & sP25) | ~ sP126), inference(usedef, [], [e136])).
fof(e136, plain, (sP126 <=> (sP49 & sP48 & sP47 & sP46 & sP45 & sP44 & sP43 & sP42 & sP41 & sP40 & sP39 & sP38 & sP37 & sP36 & sP35 & sP34 & sP33 & sP32 & sP31 & sP30 & sP29 & sP28 & sP27 & sP26 & sP25)), introduced(predicate_definition_introduction, [new_symbols(naming, [sP126])])).
fof(f137, plain, ((sP74 & sP73 & sP72 & sP71 & sP70 & sP69 & sP68 & sP67 & sP66 & sP65 & sP64 & sP63 & sP62 & sP61 & sP60 & sP59 & sP58 & sP57 & sP56 & sP55 & sP54 & sP53 & sP52 & sP51 & sP50) | ~ sP127), inference(usedef, [], [e137])).
fof(e137, plain, (sP127 <=> (sP74 & sP73 & sP72 & sP71 & sP70 & sP69 & sP68 & sP67 & sP66 & sP65 & sP64 & sP63 & sP62 & sP61 & sP60 & sP59 & sP58 & sP57 & sP56 & sP55 & sP54 & sP53 & sP52 & sP51 & sP50)), introduced(predicate_definition_introduction, [new_symbols(naming, [sP127])])).
fof(f138, plain, ((sP99 & sP98 & sP97 & sP96 & sP95 & sP94 & sP93 & sP92 & sP91 & sP90 & sP89 & sP88 & sP87 & sP86 & sP85 & sP84 & sP83 & sP82 & sP81 & sP80 & sP79 & sP78 & sP77 & sP76 & sP75) | ~ sP128), inference(usedef, [], [e138])).
fof(e138, plain, (sP128 <=> (sP99 & sP98 & sP97 & sP96 & sP95 & sP94 & sP93 & sP92 & sP91 & sP90 & sP89 & sP88 & sP87 & sP86 & sP85 & sP84 & sP83 & sP82 & sP81 & sP80 & sP79 & sP78 & sP77 & sP76 & sP75)), introduced(predicate_definition_introduction, [new_symbols(naming, [sP128])])).
fof(f139, plain, ((sP124 & sP123 & sP122 & sP121 & sP120 & sP119 & sP118 & sP117 & sP116 & sP115 & sP114 & sP113 & sP112 & sP111 & sP110 & sP109 & sP108 & sP107 & sP106 & sP105 & sP104 & sP103 & sP102 & sP101 & sP100) | ~ sP129), inference(usedef, [], [e139])).
fof(e139, plain, (sP129 <=> (sP124 & sP123 & sP122 & sP121 & sP120 & sP119 & sP118 & sP117 & sP116 & sP115 & sP114 & sP113 & sP112 & sP111 & sP110 & sP109 & sP108 & sP107 & sP106 & sP105 & sP104 & sP103 & sP102 & sP101 & sP100)), introduced(predicate_definition_introduction, [new_symbols(naming, [sP129])])).
fof(f9, plain, (((((~ (e4 = unit) & (e4 = op(e4, e4))) | ~ (e4 = op(e4, e4))) & ((~ (e4 = unit) & (e3 = op(e4, e4))) | ~ (e4 = op(e4, e3))) & ((~ (e4 = unit) & (e2 = op(e4, e4))) | ~ (e4 = op(e4, e2))) & ((~ (e4 = unit) & (e1 = op(e4, e4))) | ~ (e4 = op(e4, e1))) & ((~ (e4 = unit) & (e0 = op(e4, e4))) | ~ (e4 = op(e4, e0))) & ((~ (e4 = unit) & (e4 = op(e3, e4))) | ~ (e4 = op(e3, e4))) & ((~ (e4 = unit) & (e3 = op(e3, e4))) | ~ (e4 = op(e3, e3))) & ((~ (e4 = unit) & (e2 = op(e3, e4))) | ~ (e4 = op(e3, e2))) & ((~ (e4 = unit) & (e1 = op(e3, e4))) | ~ (e4 = op(e3, e1))) & ((~ (e4 = unit) & (e0 = op(e3, e4))) | ~ (e4 = op(e3, e0))) & ((~ (e4 = unit) & (e4 = op(e2, e4))) | ~ (e4 = op(e2, e4))) & ((~ (e4 = unit) & (e3 = op(e2, e4))) | ~ (e4 = op(e2, e3))) & ((~ (e4 = unit) & (e2 = op(e2, e4))) | ~ (e4 = op(e2, e2))) & ((~ (e4 = unit) & (e1 = op(e2, e4))) | ~ (e4 = op(e2, e1))) & ((~ (e4 = unit) & (e0 = op(e2, e4))) | ~ (e4 = op(e2, e0))) & ((~ (e4 = unit) & (e4 = op(e1, e4))) | ~ (e4 = op(e1, e4))) & ((~ (e4 = unit) & (e3 = op(e1, e4))) | ~ (e4 = op(e1, e3))) & ((~ (e4 = unit) & (e2 = op(e1, e4))) | ~ (e4 = op(e1, e2))) & ((~ (e4 = unit) & (e1 = op(e1, e4))) | ~ (e4 = op(e1, e1))) & ((~ (e4 = unit) & (e0 = op(e1, e4))) | ~ (e4 = op(e1, e0))) & ((~ (e4 = unit) & (e4 = op(e0, e4))) | ~ (e4 = op(e0, e4))) & ((~ (e4 = unit) & (e3 = op(e0, e4))) | ~ (e4 = op(e0, e3))) & ((~ (e4 = unit) & (e2 = op(e0, e4))) | ~ (e4 = op(e0, e2))) & ((~ (e4 = unit) & (e1 = op(e0, e4))) | ~ (e4 = op(e0, e1))) & ((~ (e4 = unit) & (e0 = op(e0, e4))) | ~ (op(e0, e0) = e4))) | (((~ (e3 = unit) & (e4 = op(e4, e3))) | ~ (e3 = op(e4, e4))) & ((~ (e3 = unit) & (e3 = op(e4, e3))) | ~ (e3 = op(e4, e3))) & ((~ (e3 = unit) & (e2 = op(e4, e3))) | ~ (e3 = op(e4, e2))) & ((~ (e3 = unit) & (e1 = op(e4, e3))) | ~ (e3 = op(e4, e1))) & ((~ (e3 = unit) & (e0 = op(e4, e3))) | ~ (e3 = op(e4, e0))) & ((~ (e3 = unit) & (e4 = op(e3, e3))) | ~ (e3 = op(e3, e4))) & ((~ (e3 = unit) & (e3 = op(e3, e3))) | ~ (e3 = op(e3, e3))) & ((~ (e3 = unit) & (e2 = op(e3, e3))) | ~ (e3 = op(e3, e2))) & ((~ (e3 = unit) & (e1 = op(e3, e3))) | ~ (e3 = op(e3, e1))) & ((~ (e3 = unit) & (e0 = op(e3, e3))) | ~ (e3 = op(e3, e0))) & ((~ (e3 = unit) & (e4 = op(e2, e3))) | ~ (e3 = op(e2, e4))) & ((~ (e3 = unit) & (e3 = op(e2, e3))) | ~ (e3 = op(e2, e3))) & ((~ (e3 = unit) & (e2 = op(e2, e3))) | ~ (e3 = op(e2, e2))) & ((~ (e3 = unit) & (e1 = op(e2, e3))) | ~ (e3 = op(e2, e1))) & ((~ (e3 = unit) & (e0 = op(e2, e3))) | ~ (e3 = op(e2, e0))) & ((~ (e3 = unit) & (e4 = op(e1, e3))) | ~ (e3 = op(e1, e4))) & ((~ (e3 = unit) & (e3 = op(e1, e3))) | ~ (e3 = op(e1, e3))) & ((~ (e3 = unit) & (e2 = op(e1, e3))) | ~ (e3 = op(e1, e2))) & ((~ (e3 = unit) & (e1 = op(e1, e3))) | ~ (e3 = op(e1, e1))) & ((~ (e3 = unit) & (e0 = op(e1, e3))) | ~ (e3 = op(e1, e0))) & ((~ (e3 = unit) & (e4 = op(e0, e3))) | ~ (e3 = op(e0, e4))) & ((~ (e3 = unit) & (e3 = op(e0, e3))) | ~ (e3 = op(e0, e3))) & ((~ (e3 = unit) & (e2 = op(e0, e3))) | ~ (e3 = op(e0, e2))) & ((~ (e3 = unit) & (e1 = op(e0, e3))) | ~ (e3 = op(e0, e1))) & ((~ (e3 = unit) & (e0 = op(e0, e3))) | ~ (op(e0, e0) = e3))) | (((~ (e2 = unit) & (e4 = op(e4, e2))) | ~ (e2 = op(e4, e4))) & ((~ (e2 = unit) & (e3 = op(e4, e2))) | ~ (e2 = op(e4, e3))) & ((~ (e2 = unit) & (e2 = op(e4, e2))) | ~ (e2 = op(e4, e2))) & ((~ (e2 = unit) & (e1 = op(e4, e2))) | ~ (e2 = op(e4, e1))) & ((~ (e2 = unit) & (e0 = op(e4, e2))) | ~ (e2 = op(e4, e0))) & ((~ (e2 = unit) & (e4 = op(e3, e2))) | ~ (e2 = op(e3, e4))) & ((~ (e2 = unit) & (e3 = op(e3, e2))) | ~ (e2 = op(e3, e3))) & ((~ (e2 = unit) & (e2 = op(e3, e2))) | ~ (e2 = op(e3, e2))) & ((~ (e2 = unit) & (e1 = op(e3, e2))) | ~ (e2 = op(e3, e1))) & ((~ (e2 = unit) & (e0 = op(e3, e2))) | ~ (e2 = op(e3, e0))) & ((~ (e2 = unit) & (e4 = op(e2, e2))) | ~ (e2 = op(e2, e4))) & ((~ (e2 = unit) & (e3 = op(e2, e2))) | ~ (e2 = op(e2, e3))) & ((~ (e2 = unit) & (e2 = op(e2, e2))) | ~ (e2 = op(e2, e2))) & ((~ (e2 = unit) & (e1 = op(e2, e2))) | ~ (e2 = op(e2, e1))) & ((~ (e2 = unit) & (e0 = op(e2, e2))) | ~ (e2 = op(e2, e0))) & ((~ (e2 = unit) & (e4 = op(e1, e2))) | ~ (e2 = op(e1, e4))) & ((~ (e2 = unit) & (e3 = op(e1, e2))) | ~ (e2 = op(e1, e3))) & ((~ (e2 = unit) & (e2 = op(e1, e2))) | ~ (e2 = op(e1, e2))) & ((~ (e2 = unit) & (e1 = op(e1, e2))) | ~ (e2 = op(e1, e1))) & ((~ (e2 = unit) & (e0 = op(e1, e2))) | ~ (e2 = op(e1, e0))) & ((~ (e2 = unit) & (e4 = op(e0, e2))) | ~ (e2 = op(e0, e4))) & ((~ (e2 = unit) & (e3 = op(e0, e2))) | ~ (e2 = op(e0, e3))) & ((~ (e2 = unit) & (e2 = op(e0, e2))) | ~ (e2 = op(e0, e2))) & ((~ (e2 = unit) & (e1 = op(e0, e2))) | ~ (e2 = op(e0, e1))) & ((~ (e2 = unit) & (e0 = op(e0, e2))) | ~ (op(e0, e0) = e2))) | (((~ (e1 = unit) & (e4 = op(e4, e1))) | ~ (e1 = op(e4, e4))) & ((~ (e1 = unit) & (e3 = op(e4, e1))) | ~ (e1 = op(e4, e3))) & ((~ (e1 = unit) & (e2 = op(e4, e1))) | ~ (e1 = op(e4, e2))) & ((~ (e1 = unit) & (e1 = op(e4, e1))) | ~ (e1 = op(e4, e1))) & ((~ (e1 = unit) & (e0 = op(e4, e1))) | ~ (e1 = op(e4, e0))) & ((~ (e1 = unit) & (e4 = op(e3, e1))) | ~ (e1 = op(e3, e4))) & ((~ (e1 = unit) & (e3 = op(e3, e1))) | ~ (e1 = op(e3, e3))) & ((~ (e1 = unit) & (e2 = op(e3, e1))) | ~ (e1 = op(e3, e2))) & ((~ (e1 = unit) & (e1 = op(e3, e1))) | ~ (e1 = op(e3, e1))) & ((~ (e1 = unit) & (e0 = op(e3, e1))) | ~ (e1 = op(e3, e0))) & ((~ (e1 = unit) & (e4 = op(e2, e1))) | ~ (e1 = op(e2, e4))) & ((~ (e1 = unit) & (e3 = op(e2, e1))) | ~ (e1 = op(e2, e3))) & ((~ (e1 = unit) & (e2 = op(e2, e1))) | ~ (e1 = op(e2, e2))) & ((~ (e1 = unit) & (e1 = op(e2, e1))) | ~ (e1 = op(e2, e1))) & ((~ (e1 = unit) & (e0 = op(e2, e1))) | ~ (e1 = op(e2, e0))) & ((~ (e1 = unit) & (e4 = op(e1, e1))) | ~ (e1 = op(e1, e4))) & ((~ (e1 = unit) & (e3 = op(e1, e1))) | ~ (e1 = op(e1, e3))) & ((~ (e1 = unit) & (e2 = op(e1, e1))) | ~ (e1 = op(e1, e2))) & ((~ (e1 = unit) & (e1 = op(e1, e1))) | ~ (e1 = op(e1, e1))) & ((~ (e1 = unit) & (e0 = op(e1, e1))) | ~ (e1 = op(e1, e0))) & ((~ (e1 = unit) & (e4 = op(e0, e1))) | ~ (e1 = op(e0, e4))) & ((~ (e1 = unit) & (e3 = op(e0, e1))) | ~ (e1 = op(e0, e3))) & ((~ (e1 = unit) & (e2 = op(e0, e1))) | ~ (e1 = op(e0, e2))) & ((~ (e1 = unit) & (e1 = op(e0, e1))) | ~ (e1 = op(e0, e1))) & ((~ (e1 = unit) & (e0 = op(e0, e1))) | ~ (op(e0, e0) = e1))) | (((~ (e0 = unit) & (e4 = op(e4, e0))) | ~ (e0 = op(e4, e4))) & ((~ (e0 = unit) & (e3 = op(e4, e0))) | ~ (e0 = op(e4, e3))) & ((~ (e0 = unit) & (e2 = op(e4, e0))) | ~ (e0 = op(e4, e2))) & ((~ (e0 = unit) & (e1 = op(e4, e0))) | ~ (e0 = op(e4, e1))) & ((~ (e0 = unit) & (e0 = op(e4, e0))) | ~ (e0 = op(e4, e0))) & ((~ (e0 = unit) & (e4 = op(e3, e0))) | ~ (e0 = op(e3, e4))) & ((~ (e0 = unit) & (e3 = op(e3, e0))) | ~ (e0 = op(e3, e3))) & ((~ (e0 = unit) & (e2 = op(e3, e0))) | ~ (e0 = op(e3, e2))) & ((~ (e0 = unit) & (e1 = op(e3, e0))) | ~ (e0 = op(e3, e1))) & ((~ (e0 = unit) & (e0 = op(e3, e0))) | ~ (e0 = op(e3, e0))) & ((~ (e0 = unit) & (e4 = op(e2, e0))) | ~ (e0 = op(e2, e4))) & ((~ (e0 = unit) & (e3 = op(e2, e0))) | ~ (e0 = op(e2, e3))) & ((~ (e0 = unit) & (e2 = op(e2, e0))) | ~ (e0 = op(e2, e2))) & ((~ (e0 = unit) & (e1 = op(e2, e0))) | ~ (e0 = op(e2, e1))) & ((~ (e0 = unit) & (e0 = op(e2, e0))) | ~ (e0 = op(e2, e0))) & ((~ (e0 = unit) & (e4 = op(e1, e0))) | ~ (e0 = op(e1, e4))) & ((~ (e0 = unit) & (e3 = op(e1, e0))) | ~ (e0 = op(e1, e3))) & ((~ (e0 = unit) & (e2 = op(e1, e0))) | ~ (e0 = op(e1, e2))) & ((~ (e0 = unit) & (e1 = op(e1, e0))) | ~ (e0 = op(e1, e1))) & ((~ (e0 = unit) & (e0 = op(e1, e0))) | ~ (e0 = op(e1, e0))) & ((~ (e0 = unit) & (op(e0, e0) = e4)) | ~ (e0 = op(e0, e4))) & ((~ (e0 = unit) & (op(e0, e0) = e3)) | ~ (e0 = op(e0, e3))) & ((~ (e0 = unit) & (op(e0, e0) = e2)) | ~ (e0 = op(e0, e2))) & ((~ (e0 = unit) & (op(e0, e0) = e1)) | ~ (e0 = op(e0, e1))) & ((~ (e0 = unit) & (e0 = op(e0, e0))) | ~ (e0 = op(e0, e0))))) & ((e4 = op(e4, e4)) | (e4 = op(e3, e3)) | (e4 = op(e2, e2)) | (e4 = op(e1, e1)) | (op(e0, e0) = e4)) & ((e3 = op(e4, e4)) | (e3 = op(e3, e3)) | (e3 = op(e2, e2)) | (e3 = op(e1, e1)) | (op(e0, e0) = e3)) & ((e2 = op(e4, e4)) | (e2 = op(e3, e3)) | (e2 = op(e2, e2)) | (e2 = op(e1, e1)) | (op(e0, e0) = e2)) & ((e1 = op(e4, e4)) | (e1 = op(e3, e3)) | (e1 = op(e2, e2)) | (e1 = op(e1, e1)) | (op(e0, e0) = e1)) & ((e0 = op(e4, e4)) | (e0 = op(e3, e3)) | (e0 = op(e2, e2)) | (e0 = op(e1, e1)) | (e0 = op(e0, e0)))), inference(flattening, [], [f8])).
fof(f8, plain, ~ ~ (((((~ (e4 = unit) & (e4 = op(e4, e4))) | ~ (e4 = op(e4, e4))) & ((~ (e4 = unit) & (e3 = op(e4, e4))) | ~ (e4 = op(e4, e3))) & ((~ (e4 = unit) & (e2 = op(e4, e4))) | ~ (e4 = op(e4, e2))) & ((~ (e4 = unit) & (e1 = op(e4, e4))) | ~ (e4 = op(e4, e1))) & ((~ (e4 = unit) & (e0 = op(e4, e4))) | ~ (e4 = op(e4, e0))) & ((~ (e4 = unit) & (e4 = op(e3, e4))) | ~ (e4 = op(e3, e4))) & ((~ (e4 = unit) & (e3 = op(e3, e4))) | ~ (e4 = op(e3, e3))) & ((~ (e4 = unit) & (e2 = op(e3, e4))) | ~ (e4 = op(e3, e2))) & ((~ (e4 = unit) & (e1 = op(e3, e4))) | ~ (e4 = op(e3, e1))) & ((~ (e4 = unit) & (e0 = op(e3, e4))) | ~ (e4 = op(e3, e0))) & ((~ (e4 = unit) & (e4 = op(e2, e4))) | ~ (e4 = op(e2, e4))) & ((~ (e4 = unit) & (e3 = op(e2, e4))) | ~ (e4 = op(e2, e3))) & ((~ (e4 = unit) & (e2 = op(e2, e4))) | ~ (e4 = op(e2, e2))) & ((~ (e4 = unit) & (e1 = op(e2, e4))) | ~ (e4 = op(e2, e1))) & ((~ (e4 = unit) & (e0 = op(e2, e4))) | ~ (e4 = op(e2, e0))) & ((~ (e4 = unit) & (e4 = op(e1, e4))) | ~ (e4 = op(e1, e4))) & ((~ (e4 = unit) & (e3 = op(e1, e4))) | ~ (e4 = op(e1, e3))) & ((~ (e4 = unit) & (e2 = op(e1, e4))) | ~ (e4 = op(e1, e2))) & ((~ (e4 = unit) & (e1 = op(e1, e4))) | ~ (e4 = op(e1, e1))) & ((~ (e4 = unit) & (e0 = op(e1, e4))) | ~ (e4 = op(e1, e0))) & ((~ (e4 = unit) & (e4 = op(e0, e4))) | ~ (e4 = op(e0, e4))) & ((~ (e4 = unit) & (e3 = op(e0, e4))) | ~ (e4 = op(e0, e3))) & ((~ (e4 = unit) & (e2 = op(e0, e4))) | ~ (e4 = op(e0, e2))) & ((~ (e4 = unit) & (e1 = op(e0, e4))) | ~ (e4 = op(e0, e1))) & ((~ (e4 = unit) & (e0 = op(e0, e4))) | ~ (op(e0, e0) = e4))) | (((~ (e3 = unit) & (e4 = op(e4, e3))) | ~ (e3 = op(e4, e4))) & ((~ (e3 = unit) & (e3 = op(e4, e3))) | ~ (e3 = op(e4, e3))) & ((~ (e3 = unit) & (e2 = op(e4, e3))) | ~ (e3 = op(e4, e2))) & ((~ (e3 = unit) & (e1 = op(e4, e3))) | ~ (e3 = op(e4, e1))) & ((~ (e3 = unit) & (e0 = op(e4, e3))) | ~ (e3 = op(e4, e0))) & ((~ (e3 = unit) & (e4 = op(e3, e3))) | ~ (e3 = op(e3, e4))) & ((~ (e3 = unit) & (e3 = op(e3, e3))) | ~ (e3 = op(e3, e3))) & ((~ (e3 = unit) & (e2 = op(e3, e3))) | ~ (e3 = op(e3, e2))) & ((~ (e3 = unit) & (e1 = op(e3, e3))) | ~ (e3 = op(e3, e1))) & ((~ (e3 = unit) & (e0 = op(e3, e3))) | ~ (e3 = op(e3, e0))) & ((~ (e3 = unit) & (e4 = op(e2, e3))) | ~ (e3 = op(e2, e4))) & ((~ (e3 = unit) & (e3 = op(e2, e3))) | ~ (e3 = op(e2, e3))) & ((~ (e3 = unit) & (e2 = op(e2, e3))) | ~ (e3 = op(e2, e2))) & ((~ (e3 = unit) & (e1 = op(e2, e3))) | ~ (e3 = op(e2, e1))) & ((~ (e3 = unit) & (e0 = op(e2, e3))) | ~ (e3 = op(e2, e0))) & ((~ (e3 = unit) & (e4 = op(e1, e3))) | ~ (e3 = op(e1, e4))) & ((~ (e3 = unit) & (e3 = op(e1, e3))) | ~ (e3 = op(e1, e3))) & ((~ (e3 = unit) & (e2 = op(e1, e3))) | ~ (e3 = op(e1, e2))) & ((~ (e3 = unit) & (e1 = op(e1, e3))) | ~ (e3 = op(e1, e1))) & ((~ (e3 = unit) & (e0 = op(e1, e3))) | ~ (e3 = op(e1, e0))) & ((~ (e3 = unit) & (e4 = op(e0, e3))) | ~ (e3 = op(e0, e4))) & ((~ (e3 = unit) & (e3 = op(e0, e3))) | ~ (e3 = op(e0, e3))) & ((~ (e3 = unit) & (e2 = op(e0, e3))) | ~ (e3 = op(e0, e2))) & ((~ (e3 = unit) & (e1 = op(e0, e3))) | ~ (e3 = op(e0, e1))) & ((~ (e3 = unit) & (e0 = op(e0, e3))) | ~ (op(e0, e0) = e3))) | (((~ (e2 = unit) & (e4 = op(e4, e2))) | ~ (e2 = op(e4, e4))) & ((~ (e2 = unit) & (e3 = op(e4, e2))) | ~ (e2 = op(e4, e3))) & ((~ (e2 = unit) & (e2 = op(e4, e2))) | ~ (e2 = op(e4, e2))) & ((~ (e2 = unit) & (e1 = op(e4, e2))) | ~ (e2 = op(e4, e1))) & ((~ (e2 = unit) & (e0 = op(e4, e2))) | ~ (e2 = op(e4, e0))) & ((~ (e2 = unit) & (e4 = op(e3, e2))) | ~ (e2 = op(e3, e4))) & ((~ (e2 = unit) & (e3 = op(e3, e2))) | ~ (e2 = op(e3, e3))) & ((~ (e2 = unit) & (e2 = op(e3, e2))) | ~ (e2 = op(e3, e2))) & ((~ (e2 = unit) & (e1 = op(e3, e2))) | ~ (e2 = op(e3, e1))) & ((~ (e2 = unit) & (e0 = op(e3, e2))) | ~ (e2 = op(e3, e0))) & ((~ (e2 = unit) & (e4 = op(e2, e2))) | ~ (e2 = op(e2, e4))) & ((~ (e2 = unit) & (e3 = op(e2, e2))) | ~ (e2 = op(e2, e3))) & ((~ (e2 = unit) & (e2 = op(e2, e2))) | ~ (e2 = op(e2, e2))) & ((~ (e2 = unit) & (e1 = op(e2, e2))) | ~ (e2 = op(e2, e1))) & ((~ (e2 = unit) & (e0 = op(e2, e2))) | ~ (e2 = op(e2, e0))) & ((~ (e2 = unit) & (e4 = op(e1, e2))) | ~ (e2 = op(e1, e4))) & ((~ (e2 = unit) & (e3 = op(e1, e2))) | ~ (e2 = op(e1, e3))) & ((~ (e2 = unit) & (e2 = op(e1, e2))) | ~ (e2 = op(e1, e2))) & ((~ (e2 = unit) & (e1 = op(e1, e2))) | ~ (e2 = op(e1, e1))) & ((~ (e2 = unit) & (e0 = op(e1, e2))) | ~ (e2 = op(e1, e0))) & ((~ (e2 = unit) & (e4 = op(e0, e2))) | ~ (e2 = op(e0, e4))) & ((~ (e2 = unit) & (e3 = op(e0, e2))) | ~ (e2 = op(e0, e3))) & ((~ (e2 = unit) & (e2 = op(e0, e2))) | ~ (e2 = op(e0, e2))) & ((~ (e2 = unit) & (e1 = op(e0, e2))) | ~ (e2 = op(e0, e1))) & ((~ (e2 = unit) & (e0 = op(e0, e2))) | ~ (op(e0, e0) = e2))) | (((~ (e1 = unit) & (e4 = op(e4, e1))) | ~ (e1 = op(e4, e4))) & ((~ (e1 = unit) & (e3 = op(e4, e1))) | ~ (e1 = op(e4, e3))) & ((~ (e1 = unit) & (e2 = op(e4, e1))) | ~ (e1 = op(e4, e2))) & ((~ (e1 = unit) & (e1 = op(e4, e1))) | ~ (e1 = op(e4, e1))) & ((~ (e1 = unit) & (e0 = op(e4, e1))) | ~ (e1 = op(e4, e0))) & ((~ (e1 = unit) & (e4 = op(e3, e1))) | ~ (e1 = op(e3, e4))) & ((~ (e1 = unit) & (e3 = op(e3, e1))) | ~ (e1 = op(e3, e3))) & ((~ (e1 = unit) & (e2 = op(e3, e1))) | ~ (e1 = op(e3, e2))) & ((~ (e1 = unit) & (e1 = op(e3, e1))) | ~ (e1 = op(e3, e1))) & ((~ (e1 = unit) & (e0 = op(e3, e1))) | ~ (e1 = op(e3, e0))) & ((~ (e1 = unit) & (e4 = op(e2, e1))) | ~ (e1 = op(e2, e4))) & ((~ (e1 = unit) & (e3 = op(e2, e1))) | ~ (e1 = op(e2, e3))) & ((~ (e1 = unit) & (e2 = op(e2, e1))) | ~ (e1 = op(e2, e2))) & ((~ (e1 = unit) & (e1 = op(e2, e1))) | ~ (e1 = op(e2, e1))) & ((~ (e1 = unit) & (e0 = op(e2, e1))) | ~ (e1 = op(e2, e0))) & ((~ (e1 = unit) & (e4 = op(e1, e1))) | ~ (e1 = op(e1, e4))) & ((~ (e1 = unit) & (e3 = op(e1, e1))) | ~ (e1 = op(e1, e3))) & ((~ (e1 = unit) & (e2 = op(e1, e1))) | ~ (e1 = op(e1, e2))) & ((~ (e1 = unit) & (e1 = op(e1, e1))) | ~ (e1 = op(e1, e1))) & ((~ (e1 = unit) & (e0 = op(e1, e1))) | ~ (e1 = op(e1, e0))) & ((~ (e1 = unit) & (e4 = op(e0, e1))) | ~ (e1 = op(e0, e4))) & ((~ (e1 = unit) & (e3 = op(e0, e1))) | ~ (e1 = op(e0, e3))) & ((~ (e1 = unit) & (e2 = op(e0, e1))) | ~ (e1 = op(e0, e2))) & ((~ (e1 = unit) & (e1 = op(e0, e1))) | ~ (e1 = op(e0, e1))) & ((~ (e1 = unit) & (e0 = op(e0, e1))) | ~ (op(e0, e0) = e1))) | (((~ (e0 = unit) & (e4 = op(e4, e0))) | ~ (e0 = op(e4, e4))) & ((~ (e0 = unit) & (e3 = op(e4, e0))) | ~ (e0 = op(e4, e3))) & ((~ (e0 = unit) & (e2 = op(e4, e0))) | ~ (e0 = op(e4, e2))) & ((~ (e0 = unit) & (e1 = op(e4, e0))) | ~ (e0 = op(e4, e1))) & ((~ (e0 = unit) & (e0 = op(e4, e0))) | ~ (e0 = op(e4, e0))) & ((~ (e0 = unit) & (e4 = op(e3, e0))) | ~ (e0 = op(e3, e4))) & ((~ (e0 = unit) & (e3 = op(e3, e0))) | ~ (e0 = op(e3, e3))) & ((~ (e0 = unit) & (e2 = op(e3, e0))) | ~ (e0 = op(e3, e2))) & ((~ (e0 = unit) & (e1 = op(e3, e0))) | ~ (e0 = op(e3, e1))) & ((~ (e0 = unit) & (e0 = op(e3, e0))) | ~ (e0 = op(e3, e0))) & ((~ (e0 = unit) & (e4 = op(e2, e0))) | ~ (e0 = op(e2, e4))) & ((~ (e0 = unit) & (e3 = op(e2, e0))) | ~ (e0 = op(e2, e3))) & ((~ (e0 = unit) & (e2 = op(e2, e0))) | ~ (e0 = op(e2, e2))) & ((~ (e0 = unit) & (e1 = op(e2, e0))) | ~ (e0 = op(e2, e1))) & ((~ (e0 = unit) & (e0 = op(e2, e0))) | ~ (e0 = op(e2, e0))) & ((~ (e0 = unit) & (e4 = op(e1, e0))) | ~ (e0 = op(e1, e4))) & ((~ (e0 = unit) & (e3 = op(e1, e0))) | ~ (e0 = op(e1, e3))) & ((~ (e0 = unit) & (e2 = op(e1, e0))) | ~ (e0 = op(e1, e2))) & ((~ (e0 = unit) & (e1 = op(e1, e0))) | ~ (e0 = op(e1, e1))) & ((~ (e0 = unit) & (e0 = op(e1, e0))) | ~ (e0 = op(e1, e0))) & ((~ (e0 = unit) & (op(e0, e0) = e4)) | ~ (e0 = op(e0, e4))) & ((~ (e0 = unit) & (op(e0, e0) = e3)) | ~ (e0 = op(e0, e3))) & ((~ (e0 = unit) & (op(e0, e0) = e2)) | ~ (e0 = op(e0, e2))) & ((~ (e0 = unit) & (op(e0, e0) = e1)) | ~ (e0 = op(e0, e1))) & ((~ (e0 = unit) & (e0 = op(e0, e0))) | ~ (e0 = op(e0, e0))))) & ((e4 = op(e4, e4)) | (e4 = op(e3, e3)) | (e4 = op(e2, e2)) | (e4 = op(e1, e1)) | (op(e0, e0) = e4)) & ((e3 = op(e4, e4)) | (e3 = op(e3, e3)) | (e3 = op(e2, e2)) | (e3 = op(e1, e1)) | (op(e0, e0) = e3)) & ((e2 = op(e4, e4)) | (e2 = op(e3, e3)) | (e2 = op(e2, e2)) | (e2 = op(e1, e1)) | (op(e0, e0) = e2)) & ((e1 = op(e4, e4)) | (e1 = op(e3, e3)) | (e1 = op(e2, e2)) | (e1 = op(e1, e1)) | (op(e0, e0) = e1)) & ((e0 = op(e4, e4)) | (e0 = op(e3, e3)) | (e0 = op(e2, e2)) | (e0 = op(e1, e1)) | (e0 = op(e0, e0)))), inference(negated_conjecture, [], [f7])).
fof(f7, plain, ~ ~ (((((~ (e4 = unit) & (e4 = op(e4, e4))) | ~ (e4 = op(e4, e4))) & ((~ (e4 = unit) & (e3 = op(e4, e4))) | ~ (e4 = op(e4, e3))) & ((~ (e4 = unit) & (e2 = op(e4, e4))) | ~ (e4 = op(e4, e2))) & ((~ (e4 = unit) & (e1 = op(e4, e4))) | ~ (e4 = op(e4, e1))) & ((~ (e4 = unit) & (e0 = op(e4, e4))) | ~ (e4 = op(e4, e0))) & ((~ (e4 = unit) & (e4 = op(e3, e4))) | ~ (e4 = op(e3, e4))) & ((~ (e4 = unit) & (e3 = op(e3, e4))) | ~ (e4 = op(e3, e3))) & ((~ (e4 = unit) & (e2 = op(e3, e4))) | ~ (e4 = op(e3, e2))) & ((~ (e4 = unit) & (e1 = op(e3, e4))) | ~ (e4 = op(e3, e1))) & ((~ (e4 = unit) & (e0 = op(e3, e4))) | ~ (e4 = op(e3, e0))) & ((~ (e4 = unit) & (e4 = op(e2, e4))) | ~ (e4 = op(e2, e4))) & ((~ (e4 = unit) & (e3 = op(e2, e4))) | ~ (e4 = op(e2, e3))) & ((~ (e4 = unit) & (e2 = op(e2, e4))) | ~ (e4 = op(e2, e2))) & ((~ (e4 = unit) & (e1 = op(e2, e4))) | ~ (e4 = op(e2, e1))) & ((~ (e4 = unit) & (e0 = op(e2, e4))) | ~ (e4 = op(e2, e0))) & ((~ (e4 = unit) & (e4 = op(e1, e4))) | ~ (e4 = op(e1, e4))) & ((~ (e4 = unit) & (e3 = op(e1, e4))) | ~ (e4 = op(e1, e3))) & ((~ (e4 = unit) & (e2 = op(e1, e4))) | ~ (e4 = op(e1, e2))) & ((~ (e4 = unit) & (e1 = op(e1, e4))) | ~ (e4 = op(e1, e1))) & ((~ (e4 = unit) & (e0 = op(e1, e4))) | ~ (e4 = op(e1, e0))) & ((~ (e4 = unit) & (e4 = op(e0, e4))) | ~ (e4 = op(e0, e4))) & ((~ (e4 = unit) & (e3 = op(e0, e4))) | ~ (e4 = op(e0, e3))) & ((~ (e4 = unit) & (e2 = op(e0, e4))) | ~ (e4 = op(e0, e2))) & ((~ (e4 = unit) & (e1 = op(e0, e4))) | ~ (e4 = op(e0, e1))) & ((~ (e4 = unit) & (e0 = op(e0, e4))) | ~ (op(e0, e0) = e4))) | (((~ (e3 = unit) & (e4 = op(e4, e3))) | ~ (e3 = op(e4, e4))) & ((~ (e3 = unit) & (e3 = op(e4, e3))) | ~ (e3 = op(e4, e3))) & ((~ (e3 = unit) & (e2 = op(e4, e3))) | ~ (e3 = op(e4, e2))) & ((~ (e3 = unit) & (e1 = op(e4, e3))) | ~ (e3 = op(e4, e1))) & ((~ (e3 = unit) & (e0 = op(e4, e3))) | ~ (e3 = op(e4, e0))) & ((~ (e3 = unit) & (e4 = op(e3, e3))) | ~ (e3 = op(e3, e4))) & ((~ (e3 = unit) & (e3 = op(e3, e3))) | ~ (e3 = op(e3, e3))) & ((~ (e3 = unit) & (e2 = op(e3, e3))) | ~ (e3 = op(e3, e2))) & ((~ (e3 = unit) & (e1 = op(e3, e3))) | ~ (e3 = op(e3, e1))) & ((~ (e3 = unit) & (e0 = op(e3, e3))) | ~ (e3 = op(e3, e0))) & ((~ (e3 = unit) & (e4 = op(e2, e3))) | ~ (e3 = op(e2, e4))) & ((~ (e3 = unit) & (e3 = op(e2, e3))) | ~ (e3 = op(e2, e3))) & ((~ (e3 = unit) & (e2 = op(e2, e3))) | ~ (e3 = op(e2, e2))) & ((~ (e3 = unit) & (e1 = op(e2, e3))) | ~ (e3 = op(e2, e1))) & ((~ (e3 = unit) & (e0 = op(e2, e3))) | ~ (e3 = op(e2, e0))) & ((~ (e3 = unit) & (e4 = op(e1, e3))) | ~ (e3 = op(e1, e4))) & ((~ (e3 = unit) & (e3 = op(e1, e3))) | ~ (e3 = op(e1, e3))) & ((~ (e3 = unit) & (e2 = op(e1, e3))) | ~ (e3 = op(e1, e2))) & ((~ (e3 = unit) & (e1 = op(e1, e3))) | ~ (e3 = op(e1, e1))) & ((~ (e3 = unit) & (e0 = op(e1, e3))) | ~ (e3 = op(e1, e0))) & ((~ (e3 = unit) & (e4 = op(e0, e3))) | ~ (e3 = op(e0, e4))) & ((~ (e3 = unit) & (e3 = op(e0, e3))) | ~ (e3 = op(e0, e3))) & ((~ (e3 = unit) & (e2 = op(e0, e3))) | ~ (e3 = op(e0, e2))) & ((~ (e3 = unit) & (e1 = op(e0, e3))) | ~ (e3 = op(e0, e1))) & ((~ (e3 = unit) & (e0 = op(e0, e3))) | ~ (op(e0, e0) = e3))) | (((~ (e2 = unit) & (e4 = op(e4, e2))) | ~ (e2 = op(e4, e4))) & ((~ (e2 = unit) & (e3 = op(e4, e2))) | ~ (e2 = op(e4, e3))) & ((~ (e2 = unit) & (e2 = op(e4, e2))) | ~ (e2 = op(e4, e2))) & ((~ (e2 = unit) & (e1 = op(e4, e2))) | ~ (e2 = op(e4, e1))) & ((~ (e2 = unit) & (e0 = op(e4, e2))) | ~ (e2 = op(e4, e0))) & ((~ (e2 = unit) & (e4 = op(e3, e2))) | ~ (e2 = op(e3, e4))) & ((~ (e2 = unit) & (e3 = op(e3, e2))) | ~ (e2 = op(e3, e3))) & ((~ (e2 = unit) & (e2 = op(e3, e2))) | ~ (e2 = op(e3, e2))) & ((~ (e2 = unit) & (e1 = op(e3, e2))) | ~ (e2 = op(e3, e1))) & ((~ (e2 = unit) & (e0 = op(e3, e2))) | ~ (e2 = op(e3, e0))) & ((~ (e2 = unit) & (e4 = op(e2, e2))) | ~ (e2 = op(e2, e4))) & ((~ (e2 = unit) & (e3 = op(e2, e2))) | ~ (e2 = op(e2, e3))) & ((~ (e2 = unit) & (e2 = op(e2, e2))) | ~ (e2 = op(e2, e2))) & ((~ (e2 = unit) & (e1 = op(e2, e2))) | ~ (e2 = op(e2, e1))) & ((~ (e2 = unit) & (e0 = op(e2, e2))) | ~ (e2 = op(e2, e0))) & ((~ (e2 = unit) & (e4 = op(e1, e2))) | ~ (e2 = op(e1, e4))) & ((~ (e2 = unit) & (e3 = op(e1, e2))) | ~ (e2 = op(e1, e3))) & ((~ (e2 = unit) & (e2 = op(e1, e2))) | ~ (e2 = op(e1, e2))) & ((~ (e2 = unit) & (e1 = op(e1, e2))) | ~ (e2 = op(e1, e1))) & ((~ (e2 = unit) & (e0 = op(e1, e2))) | ~ (e2 = op(e1, e0))) & ((~ (e2 = unit) & (e4 = op(e0, e2))) | ~ (e2 = op(e0, e4))) & ((~ (e2 = unit) & (e3 = op(e0, e2))) | ~ (e2 = op(e0, e3))) & ((~ (e2 = unit) & (e2 = op(e0, e2))) | ~ (e2 = op(e0, e2))) & ((~ (e2 = unit) & (e1 = op(e0, e2))) | ~ (e2 = op(e0, e1))) & ((~ (e2 = unit) & (e0 = op(e0, e2))) | ~ (op(e0, e0) = e2))) | (((~ (e1 = unit) & (e4 = op(e4, e1))) | ~ (e1 = op(e4, e4))) & ((~ (e1 = unit) & (e3 = op(e4, e1))) | ~ (e1 = op(e4, e3))) & ((~ (e1 = unit) & (e2 = op(e4, e1))) | ~ (e1 = op(e4, e2))) & ((~ (e1 = unit) & (e1 = op(e4, e1))) | ~ (e1 = op(e4, e1))) & ((~ (e1 = unit) & (e0 = op(e4, e1))) | ~ (e1 = op(e4, e0))) & ((~ (e1 = unit) & (e4 = op(e3, e1))) | ~ (e1 = op(e3, e4))) & ((~ (e1 = unit) & (e3 = op(e3, e1))) | ~ (e1 = op(e3, e3))) & ((~ (e1 = unit) & (e2 = op(e3, e1))) | ~ (e1 = op(e3, e2))) & ((~ (e1 = unit) & (e1 = op(e3, e1))) | ~ (e1 = op(e3, e1))) & ((~ (e1 = unit) & (e0 = op(e3, e1))) | ~ (e1 = op(e3, e0))) & ((~ (e1 = unit) & (e4 = op(e2, e1))) | ~ (e1 = op(e2, e4))) & ((~ (e1 = unit) & (e3 = op(e2, e1))) | ~ (e1 = op(e2, e3))) & ((~ (e1 = unit) & (e2 = op(e2, e1))) | ~ (e1 = op(e2, e2))) & ((~ (e1 = unit) & (e1 = op(e2, e1))) | ~ (e1 = op(e2, e1))) & ((~ (e1 = unit) & (e0 = op(e2, e1))) | ~ (e1 = op(e2, e0))) & ((~ (e1 = unit) & (e4 = op(e1, e1))) | ~ (e1 = op(e1, e4))) & ((~ (e1 = unit) & (e3 = op(e1, e1))) | ~ (e1 = op(e1, e3))) & ((~ (e1 = unit) & (e2 = op(e1, e1))) | ~ (e1 = op(e1, e2))) & ((~ (e1 = unit) & (e1 = op(e1, e1))) | ~ (e1 = op(e1, e1))) & ((~ (e1 = unit) & (e0 = op(e1, e1))) | ~ (e1 = op(e1, e0))) & ((~ (e1 = unit) & (e4 = op(e0, e1))) | ~ (e1 = op(e0, e4))) & ((~ (e1 = unit) & (e3 = op(e0, e1))) | ~ (e1 = op(e0, e3))) & ((~ (e1 = unit) & (e2 = op(e0, e1))) | ~ (e1 = op(e0, e2))) & ((~ (e1 = unit) & (e1 = op(e0, e1))) | ~ (e1 = op(e0, e1))) & ((~ (e1 = unit) & (e0 = op(e0, e1))) | ~ (op(e0, e0) = e1))) | (((~ (e0 = unit) & (e4 = op(e4, e0))) | ~ (e0 = op(e4, e4))) & ((~ (e0 = unit) & (e3 = op(e4, e0))) | ~ (e0 = op(e4, e3))) & ((~ (e0 = unit) & (e2 = op(e4, e0))) | ~ (e0 = op(e4, e2))) & ((~ (e0 = unit) & (e1 = op(e4, e0))) | ~ (e0 = op(e4, e1))) & ((~ (e0 = unit) & (e0 = op(e4, e0))) | ~ (e0 = op(e4, e0))) & ((~ (e0 = unit) & (e4 = op(e3, e0))) | ~ (e0 = op(e3, e4))) & ((~ (e0 = unit) & (e3 = op(e3, e0))) | ~ (e0 = op(e3, e3))) & ((~ (e0 = unit) & (e2 = op(e3, e0))) | ~ (e0 = op(e3, e2))) & ((~ (e0 = unit) & (e1 = op(e3, e0))) | ~ (e0 = op(e3, e1))) & ((~ (e0 = unit) & (e0 = op(e3, e0))) | ~ (e0 = op(e3, e0))) & ((~ (e0 = unit) & (e4 = op(e2, e0))) | ~ (e0 = op(e2, e4))) & ((~ (e0 = unit) & (e3 = op(e2, e0))) | ~ (e0 = op(e2, e3))) & ((~ (e0 = unit) & (e2 = op(e2, e0))) | ~ (e0 = op(e2, e2))) & ((~ (e0 = unit) & (e1 = op(e2, e0))) | ~ (e0 = op(e2, e1))) & ((~ (e0 = unit) & (e0 = op(e2, e0))) | ~ (e0 = op(e2, e0))) & ((~ (e0 = unit) & (e4 = op(e1, e0))) | ~ (e0 = op(e1, e4))) & ((~ (e0 = unit) & (e3 = op(e1, e0))) | ~ (e0 = op(e1, e3))) & ((~ (e0 = unit) & (e2 = op(e1, e0))) | ~ (e0 = op(e1, e2))) & ((~ (e0 = unit) & (e1 = op(e1, e0))) | ~ (e0 = op(e1, e1))) & ((~ (e0 = unit) & (e0 = op(e1, e0))) | ~ (e0 = op(e1, e0))) & ((~ (e0 = unit) & (op(e0, e0) = e4)) | ~ (e0 = op(e0, e4))) & ((~ (e0 = unit) & (op(e0, e0) = e3)) | ~ (e0 = op(e0, e3))) & ((~ (e0 = unit) & (op(e0, e0) = e2)) | ~ (e0 = op(e0, e2))) & ((~ (e0 = unit) & (op(e0, e0) = e1)) | ~ (e0 = op(e0, e1))) & ((~ (e0 = unit) & (e0 = op(e0, e0))) | ~ (e0 = op(e0, e0))))) & ((e4 = op(e4, e4)) | (e4 = op(e3, e3)) | (e4 = op(e2, e2)) | (e4 = op(e1, e1)) | (op(e0, e0) = e4)) & ((e3 = op(e4, e4)) | (e3 = op(e3, e3)) | (e3 = op(e2, e2)) | (e3 = op(e1, e1)) | (op(e0, e0) = e3)) & ((e2 = op(e4, e4)) | (e2 = op(e3, e3)) | (e2 = op(e2, e2)) | (e2 = op(e1, e1)) | (op(e0, e0) = e2)) & ((e1 = op(e4, e4)) | (e1 = op(e3, e3)) | (e1 = op(e2, e2)) | (e1 = op(e1, e1)) | (op(e0, e0) = e1)) & ((e0 = op(e4, e4)) | (e0 = op(e3, e3)) | (e0 = op(e2, e2)) | (e0 = op(e1, e1)) | (e0 = op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG062+1.p', co1)).
fof(f2323, plain, (spl130_235 | spl130_209 | spl130_183 | spl130_157 | spl130_131), inference(avatar_split_clause, [], [f851, f1454, f1583, f1712, f1841, f1970])).
fof(f1970, plain, (spl130_235 <=> sP125), introduced(avatar_definition, [new_symbols(naming, [spl130_235])])).
fof(f1841, plain, (spl130_209 <=> sP126), introduced(avatar_definition, [new_symbols(naming, [spl130_209])])).
fof(f1712, plain, (spl130_183 <=> sP127), introduced(avatar_definition, [new_symbols(naming, [spl130_183])])).
fof(f1583, plain, (spl130_157 <=> sP128), introduced(avatar_definition, [new_symbols(naming, [spl130_157])])).
fof(f1454, plain, (spl130_131 <=> sP129), introduced(avatar_definition, [new_symbols(naming, [spl130_131])])).
fof(f851, plain, (sP129 | sP128 | sP127 | sP126 | sP125), inference(cnf_transformation, [], [f140])).
fof(f2309, plain, (~ spl130_253 | ~ spl130_86 | ~ spl130_126), inference(avatar_split_clause, [], [f831, f1378, f1210, f2059])).
fof(f2059, plain, (spl130_253 <=> sP7), introduced(avatar_definition, [new_symbols(naming, [spl130_253])])).
fof(f831, plain, (~ (e0 = unit) | ~ (e0 = op(e1, e2)) | ~ sP7), inference(cnf_transformation, [], [f263])).
fof(f263, plain, ((~ (e0 = unit) & (e2 = op(e1, e0))) | ~ (e0 = op(e1, e2)) | ~ sP7), inference(nnf_transformation, [], [f17])).
fof(f2268, plain, (~ spl130_229 | ~ spl130_97 | spl130_91), inference(avatar_split_clause, [], [f784, f1231, f1256, f1940])).
fof(f1940, plain, (spl130_229 <=> sP30), introduced(avatar_definition, [new_symbols(naming, [spl130_229])])).
fof(f784, plain, ((e0 = op(e1, e1)) | ~ (e1 = op(e1, e0)) | ~ sP30), inference(cnf_transformation, [], [f240])).
fof(f240, plain, ((~ (e1 = unit) & (e0 = op(e1, e1))) | ~ (e1 = op(e1, e0)) | ~ sP30), inference(nnf_transformation, [], [f40])).
fof(f2214, plain, (~ spl130_198 | ~ spl130_73 | spl130_61), inference(avatar_split_clause, [], [f724, f1105, f1155, f1786])).
fof(f1786, plain, (spl130_198 <=> sP60), introduced(avatar_definition, [new_symbols(naming, [spl130_198])])).
fof(f724, plain, ((e0 = op(e2, e2)) | ~ (e2 = op(e2, e0)) | ~ sP60), inference(cnf_transformation, [], [f210])).
fof(f210, plain, ((~ (e2 = unit) & (e0 = op(e2, e2))) | ~ (e2 = op(e2, e0)) | ~ sP60), inference(nnf_transformation, [], [f70])).
fof(f2171, plain, (~ spl130_173 | ~ spl130_79 | spl130_85), inference(avatar_split_clause, [], [f676, f1205, f1180, f1662])).
fof(f1662, plain, (spl130_173 <=> sP84), introduced(avatar_definition, [new_symbols(naming, [spl130_173])])).
fof(f1180, plain, (spl130_79 <=> (e3 = op(e1, e4))), introduced(avatar_definition, [new_symbols(naming, [spl130_79])])).
fof(f676, plain, ((e4 = op(e1, e3)) | ~ (e3 = op(e1, e4)) | ~ sP84), inference(cnf_transformation, [], [f186])).
fof(f186, plain, ((~ (e3 = unit) & (e4 = op(e1, e3))) | ~ (e3 = op(e1, e4)) | ~ sP84), inference(nnf_transformation, [], [f94])).
fof(f2106, plain, (~ spl130_136 | ~ spl130_25 | spl130_1), inference(avatar_split_clause, [], [f604, f853, f953, f1478])).
fof(f1478, plain, (spl130_136 <=> sP120), introduced(avatar_definition, [new_symbols(naming, [spl130_136])])).
fof(f604, plain, ((e0 = op(e4, e4)) | ~ (e4 = op(e4, e0)) | ~ sP120), inference(cnf_transformation, [], [f150])).
fof(f150, plain, ((~ (e4 = unit) & (e0 = op(e4, e4))) | ~ (e4 = op(e4, e0)) | ~ sP120), inference(nnf_transformation, [], [f130])).
fof(f2100, plain, (~ spl130_133 | ~ spl130_10 | spl130_4), inference(avatar_split_clause, [], [f598, f865, f890, f1463])).
fof(f1463, plain, (spl130_133 <=> sP123), introduced(avatar_definition, [new_symbols(naming, [spl130_133])])).
fof(f598, plain, ((e3 = op(e4, e4)) | ~ (e4 = op(e4, e3)) | ~ sP123), inference(cnf_transformation, [], [f147])).
fof(f147, plain, ((~ (e4 = unit) & (e3 = op(e4, e4))) | ~ (e4 = op(e4, e3)) | ~ sP123), inference(nnf_transformation, [], [f133])).
fof(f2062, plain, (~ spl130_235 | spl130_253), inference(avatar_split_clause, [], [f578, f2059, f1970])).
fof(f578, plain, (sP7 | ~ sP125), inference(cnf_transformation, [], [f145])).
fof(f145, plain, ((sP24 & sP23 & sP22 & sP21 & sP20 & sP19 & sP18 & sP17 & sP16 & sP15 & sP14 & sP13 & sP12 & sP11 & sP10 & sP9 & sP8 & sP7 & sP6 & sP5 & sP4 & sP3 & sP2 & sP1 & sP0) | ~ sP125), inference(nnf_transformation, [], [f135])).
fof(f1943, plain, (~ spl130_209 | spl130_229), inference(avatar_split_clause, [], [f551, f1940, f1841])).
fof(f551, plain, (sP30 | ~ sP126), inference(cnf_transformation, [], [f144])).
fof(f144, plain, ((sP49 & sP48 & sP47 & sP46 & sP45 & sP44 & sP43 & sP42 & sP41 & sP40 & sP39 & sP38 & sP37 & sP36 & sP35 & sP34 & sP33 & sP32 & sP31 & sP30 & sP29 & sP28 & sP27 & sP26 & sP25) | ~ sP126), inference(nnf_transformation, [], [f136])).
fof(f1789, plain, (~ spl130_183 | spl130_198), inference(avatar_split_clause, [], [f531, f1786, f1712])).
fof(f531, plain, (sP60 | ~ sP127), inference(cnf_transformation, [], [f143])).
fof(f143, plain, ((sP74 & sP73 & sP72 & sP71 & sP70 & sP69 & sP68 & sP67 & sP66 & sP65 & sP64 & sP63 & sP62 & sP61 & sP60 & sP59 & sP58 & sP57 & sP56 & sP55 & sP54 & sP53 & sP52 & sP51 & sP50) | ~ sP127), inference(nnf_transformation, [], [f137])).
fof(f1665, plain, (~ spl130_157 | spl130_173), inference(avatar_split_clause, [], [f505, f1662, f1583])).
fof(f505, plain, (sP84 | ~ sP128), inference(cnf_transformation, [], [f142])).
fof(f142, plain, ((sP99 & sP98 & sP97 & sP96 & sP95 & sP94 & sP93 & sP92 & sP91 & sP90 & sP89 & sP88 & sP87 & sP86 & sP85 & sP84 & sP83 & sP82 & sP81 & sP80 & sP79 & sP78 & sP77 & sP76 & sP75) | ~ sP128), inference(nnf_transformation, [], [f138])).
fof(f1481, plain, (~ spl130_131 | spl130_136), inference(avatar_split_clause, [], [f491, f1478, f1454])).
fof(f491, plain, (sP120 | ~ sP129), inference(cnf_transformation, [], [f141])).
fof(f141, plain, ((sP124 & sP123 & sP122 & sP121 & sP120 & sP119 & sP118 & sP117 & sP116 & sP115 & sP114 & sP113 & sP112 & sP111 & sP110 & sP109 & sP108 & sP107 & sP106 & sP105 & sP104 & sP103 & sP102 & sP101 & sP100) | ~ sP129), inference(nnf_transformation, [], [f139])).
fof(f1466, plain, (~ spl130_131 | spl130_133), inference(avatar_split_clause, [], [f494, f1463, f1454])).
fof(f494, plain, (sP123 | ~ sP129), inference(cnf_transformation, [], [f141])).
fof(f1452, plain, spl130_86, inference(avatar_split_clause, [], [f1451, f1210])).
fof(f1451, plain, (e0 = op(e1, e2)), inference(forward_demodulation, [], [f467, f468])).
fof(f467, plain, (e0 = op(e1, op(op(e1, e1), op(e1, e1)))), inference(cnf_transformation, [], [f6])).
fof(f1450, plain, spl130_64, inference(avatar_split_clause, [], [f1449, f1117])).
fof(f1449, plain, (e3 = op(e2, e2)), inference(backward_demodulation, [], [f469, f468])).
fof(f469, plain, (e3 = op(op(op(e1, e1), op(e1, e1)), op(op(e1, e1), op(e1, e1)))), inference(cnf_transformation, [], [f6])).
fof(f1448, plain, spl130_95, inference(avatar_split_clause, [], [f470, f1247])).
fof(f470, plain, (e4 = op(e1, e1)), inference(cnf_transformation, [], [f6])).
fof(f1400, plain, (spl130_104 | spl130_79 | spl130_54 | spl130_29 | spl130_4), inference(avatar_split_clause, [], [f354, f865, f970, f1075, f1180, f1285])).
fof(f354, plain, ((e3 = op(e4, e4)) | (e3 = op(e3, e4)) | (e3 = op(e2, e4)) | (e3 = op(e1, e4)) | (e3 = op(e0, e4))), inference(cnf_transformation, [], [f3])).
fof(f3, plain, (((e4 = op(e4, e4)) | (e4 = op(e3, e4)) | (e4 = op(e2, e4)) | (e4 = op(e1, e4)) | (e4 = op(e0, e4))) & ((e4 = op(e4, e4)) | (e4 = op(e4, e3)) | (e4 = op(e4, e2)) | (e4 = op(e4, e1)) | (e4 = op(e4, e0))) & ((e3 = op(e4, e4)) | (e3 = op(e3, e4)) | (e3 = op(e2, e4)) | (e3 = op(e1, e4)) | (e3 = op(e0, e4))) & ((e3 = op(e4, e4)) | (e3 = op(e4, e3)) | (e3 = op(e4, e2)) | (e3 = op(e4, e1)) | (e3 = op(e4, e0))) & ((e2 = op(e4, e4)) | (e2 = op(e3, e4)) | (e2 = op(e2, e4)) | (e2 = op(e1, e4)) | (e2 = op(e0, e4))) & ((e2 = op(e4, e4)) | (e2 = op(e4, e3)) | (e2 = op(e4, e2)) | (e2 = op(e4, e1)) | (e2 = op(e4, e0))) & ((e1 = op(e4, e4)) | (e1 = op(e3, e4)) | (e1 = op(e2, e4)) | (e1 = op(e1, e4)) | (e1 = op(e0, e4))) & ((e1 = op(e4, e4)) | (e1 = op(e4, e3)) | (e1 = op(e4, e2)) | (e1 = op(e4, e1)) | (e1 = op(e4, e0))) & ((e0 = op(e4, e4)) | (e0 = op(e3, e4)) | (e0 = op(e2, e4)) | (e0 = op(e1, e4)) | (e0 = op(e0, e4))) & ((e0 = op(e4, e4)) | (e0 = op(e4, e3)) | (e0 = op(e4, e2)) | (e0 = op(e4, e1)) | (e0 = op(e4, e0))) & ((e4 = op(e4, e3)) | (e4 = op(e3, e3)) | (e4 = op(e2, e3)) | (e4 = op(e1, e3)) | (e4 = op(e0, e3))) & ((e4 = op(e3, e4)) | (e4 = op(e3, e3)) | (e4 = op(e3, e2)) | (e4 = op(e3, e1)) | (e4 = op(e3, e0))) & ((e3 = op(e4, e3)) | (e3 = op(e3, e3)) | (e3 = op(e2, e3)) | (e3 = op(e1, e3)) | (e3 = op(e0, e3))) & ((e3 = op(e3, e4)) | (e3 = op(e3, e3)) | (e3 = op(e3, e2)) | (e3 = op(e3, e1)) | (e3 = op(e3, e0))) & ((e2 = op(e4, e3)) | (e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(e0, e3))) & ((e2 = op(e3, e4)) | (e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, e0))) & ((e1 = op(e4, e3)) | (e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(e0, e3))) & ((e1 = op(e3, e4)) | (e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, e0))) & ((e0 = op(e4, e3)) | (e0 = op(e3, e3)) | (e0 = op(e2, e3)) | (e0 = op(e1, e3)) | (e0 = op(e0, e3))) & ((e0 = op(e3, e4)) | (e0 = op(e3, e3)) | (e0 = op(e3, e2)) | (e0 = op(e3, e1)) | (e0 = op(e3, e0))) & ((e4 = op(e4, e2)) | (e4 = op(e3, e2)) | (e4 = op(e2, e2)) | (e4 = op(e1, e2)) | (e4 = op(e0, e2))) & ((e4 = op(e2, e4)) | (e4 = op(e2, e3)) | (e4 = op(e2, e2)) | (e4 = op(e2, e1)) | (e4 = op(e2, e0))) & ((e3 = op(e4, e2)) | (e3 = op(e3, e2)) | (e3 = op(e2, e2)) | (e3 = op(e1, e2)) | (e3 = op(e0, e2))) & ((e3 = op(e2, e4)) | (e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, e0))) & ((e2 = op(e4, e2)) | (e2 = op(e3, e2)) | (e2 = op(e2, e2)) | (e2 = op(e1, e2)) | (e2 = op(e0, e2))) & ((e2 = op(e2, e4)) | (e2 = op(e2, e3)) | (e2 = op(e2, e2)) | (e2 = op(e2, e1)) | (e2 = op(e2, e0))) & ((e1 = op(e4, e2)) | (e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(e0, e2))) & ((e1 = op(e2, e4)) | (e1 = op(e2, e3)) | (e1 = op(e2, e2)) | (e1 = op(e2, e1)) | (e1 = op(e2, e0))) & ((e0 = op(e4, e2)) | (e0 = op(e3, e2)) | (e0 = op(e2, e2)) | (e0 = op(e1, e2)) | (e0 = op(e0, e2))) & ((e0 = op(e2, e4)) | (e0 = op(e2, e3)) | (e0 = op(e2, e2)) | (e0 = op(e2, e1)) | (e0 = op(e2, e0))) & ((e4 = op(e4, e1)) | (e4 = op(e3, e1)) | (e4 = op(e2, e1)) | (e4 = op(e1, e1)) | (e4 = op(e0, e1))) & ((e4 = op(e1, e4)) | (e4 = op(e1, e3)) | (e4 = op(e1, e2)) | (e4 = op(e1, e1)) | (e4 = op(e1, e0))) & ((e3 = op(e4, e1)) | (e3 = op(e3, e1)) | (e3 = op(e2, e1)) | (e3 = op(e1, e1)) | (e3 = op(e0, e1))) & ((e3 = op(e1, e4)) | (e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, e0))) & ((e2 = op(e4, e1)) | (e2 = op(e3, e1)) | (e2 = op(e2, e1)) | (e2 = op(e1, e1)) | (e2 = op(e0, e1))) & ((e2 = op(e1, e4)) | (e2 = op(e1, e3)) | (e2 = op(e1, e2)) | (e2 = op(e1, e1)) | (e2 = op(e1, e0))) & ((e1 = op(e4, e1)) | (e1 = op(e3, e1)) | (e1 = op(e2, e1)) | (e1 = op(e1, e1)) | (e1 = op(e0, e1))) & ((e1 = op(e1, e4)) | (e1 = op(e1, e3)) | (e1 = op(e1, e2)) | (e1 = op(e1, e1)) | (e1 = op(e1, e0))) & ((e0 = op(e4, e1)) | (e0 = op(e3, e1)) | (e0 = op(e2, e1)) | (e0 = op(e1, e1)) | (e0 = op(e0, e1))) & ((e0 = op(e1, e4)) | (e0 = op(e1, e3)) | (e0 = op(e1, e2)) | (e0 = op(e1, e1)) | (e0 = op(e1, e0))) & ((e4 = op(e4, e0)) | (e4 = op(e3, e0)) | (e4 = op(e2, e0)) | (e4 = op(e1, e0)) | (op(e0, e0) = e4)) & ((e4 = op(e0, e4)) | (e4 = op(e0, e3)) | (e4 = op(e0, e2)) | (e4 = op(e0, e1)) | (op(e0, e0) = e4)) & ((e3 = op(e4, e0)) | (e3 = op(e3, e0)) | (e3 = op(e2, e0)) | (e3 = op(e1, e0)) | (op(e0, e0) = e3)) & ((e3 = op(e0, e4)) | (e3 = op(e0, e3)) | (e3 = op(e0, e2)) | (e3 = op(e0, e1)) | (op(e0, e0) = e3)) & ((e2 = op(e4, e0)) | (e2 = op(e3, e0)) | (e2 = op(e2, e0)) | (e2 = op(e1, e0)) | (op(e0, e0) = e2)) & ((e2 = op(e0, e4)) | (e2 = op(e0, e3)) | (e2 = op(e0, e2)) | (e2 = op(e0, e1)) | (op(e0, e0) = e2)) & ((e1 = op(e4, e0)) | (e1 = op(e3, e0)) | (e1 = op(e2, e0)) | (e1 = op(e1, e0)) | (op(e0, e0) = e1)) & ((e1 = op(e0, e4)) | (e1 = op(e0, e3)) | (e1 = op(e0, e2)) | (e1 = op(e0, e1)) | (op(e0, e0) = e1)) & ((e0 = op(e4, e0)) | (e0 = op(e3, e0)) | (e0 = op(e2, e0)) | (e0 = op(e1, e0)) | (e0 = op(e0, e0))) & ((e0 = op(e0, e4)) | (e0 = op(e0, e3)) | (e0 = op(e0, e2)) | (e0 = op(e0, e1)) | (e0 = op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG062+1.p', ax3)).
fof(f1397, plain, (spl130_126 | spl130_127 | spl130_128 | spl130_129 | spl130_130), inference(avatar_split_clause, [], [f306, f1394, f1390, f1386, f1382, f1378])).
fof(f306, plain, ((e4 = unit) | (e3 = unit) | (e2 = unit) | (e1 = unit) | (e0 = unit)), inference(cnf_transformation, [], [f2])).