fof(f6293, plain, $false, inference(avatar_sat_refutation, [], [f773, f824, f858, f943, f1045, f1181, f1254, f1282, f1285, f1286, f1289, f1442, f1460, f1481, f1817, f2173, f2236, f2237, f2239, f2240, f2345, f2397, f2510, f2572, f2894, f2937, f2939, f2951, f3048, f3061, f3093, f3175, f3218, f3281, f3303, f3334, f3383, f3420, f3421, f3544, f3563, f3757, f3981, f3982, f4017, f4161, f4212, f4535, f4706, f4768, f5165, f5314, f5342, f5375, f5376, f5377, f5378, f5383, f5418, f5436, f5727, f5728, f5729, f5771, f5854, f5899, f5934, f5968, f5998, f6030, f6063, f6097, f6129, f6158, f6190, f6219, f6248, f6268, f6287])).
fof(f6287, plain, (~ spl12_58 | ~ spl12_142 | ~ spl12_189 | spl12_249), inference(avatar_contradiction_clause, [], [f6286])).
fof(f6286, plain, ($false | (~ spl12_58 | ~ spl12_142 | ~ spl12_189 | spl12_249)), inference(subsumption_resolution, [], [f6285, f1138])).
fof(f1138, plain, ((e21 = op2(e20, e21)) | ~ spl12_142), inference(avatar_component_clause, [], [f1136])).
fof(f1136, plain, (spl12_142 <=> (e21 = op2(e20, e21))), introduced(avatar_definition, [new_symbols(naming, [spl12_142])])).
fof(f6285, plain, (~ (e21 = op2(e20, e21)) | (~ spl12_58 | ~ spl12_189 | spl12_249)), inference(forward_demodulation, [], [f6284, f1339])).
fof(f1339, plain, (e21 = h3(e11)), inference(forward_demodulation, [], [f1338, f1336])).
fof(f1336, plain, (e21 = op2(e22, h3(e13))), inference(backward_demodulation, [], [f454, f467])).
fof(f467, plain, (op2(e22, e22) = h3(e13)), inference(cnf_transformation, [], [f28])).
fof(f28, plain, ((op2(e22, e22) = h3(e13)) & (op2(e22, op2(e22, e22)) = h3(e11)) & (op2(op2(e22, op2(e22, e22)), e22) = h3(e10)) & (e22 = h3(e12))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG016+1.p', ax28)).
fof(f454, plain, (e21 = op2(e22, op2(e22, e22))), inference(cnf_transformation, [], [f25])).
fof(f25, plain, ((e23 = op2(e22, e22)) & (e21 = op2(e22, op2(e22, e22))) & (e20 = op2(op2(e22, op2(e22, e22)), e22))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG016+1.p', ax25)).
fof(f1338, plain, (h3(e11) = op2(e22, h3(e13))), inference(forward_demodulation, [], [f466, f467])).
fof(f466, plain, (op2(e22, op2(e22, e22)) = h3(e11)), inference(cnf_transformation, [], [f28])).
fof(f6284, plain, (~ (op2(e20, e21) = h3(e11)) | (~ spl12_58 | ~ spl12_189 | spl12_249)), inference(forward_demodulation, [], [f6283, f781])).
fof(f781, plain, ((e11 = op1(e10, e11)) | ~ spl12_58), inference(avatar_component_clause, [], [f779])).
fof(f779, plain, (spl12_58 <=> (e11 = op1(e10, e11))), introduced(avatar_definition, [new_symbols(naming, [spl12_58])])).
fof(f6283, plain, (~ (op2(e20, e21) = h3(op1(e10, e11))) | (~ spl12_189 | spl12_249)), inference(forward_demodulation, [], [f1816, f1479])).
fof(f1479, plain, ((e20 = h3(e10)) | ~ spl12_189), inference(avatar_component_clause, [], [f1478])).
fof(f1478, plain, (spl12_189 <=> (e20 = h3(e10))), introduced(avatar_definition, [new_symbols(naming, [spl12_189])])).
fof(f1816, plain, (~ (h3(op1(e10, e11)) = op2(h3(e10), e21)) | spl12_249), inference(avatar_component_clause, [], [f1814])).
fof(f1814, plain, (spl12_249 <=> (h3(op1(e10, e11)) = op2(h3(e10), e21))), introduced(avatar_definition, [new_symbols(naming, [spl12_249])])).
fof(f6268, plain, (~ spl12_37 | spl12_245), inference(avatar_contradiction_clause, [], [f6267])).
fof(f6267, plain, ($false | (~ spl12_37 | spl12_245)), inference(trivial_inequality_removal, [], [f6266])).
fof(f6266, plain, (~ (h3(e10) = h3(e10)) | (~ spl12_37 | spl12_245)), inference(forward_demodulation, [], [f1800, f692])).
fof(f692, plain, ((e10 = op1(e11, e12)) | ~ spl12_37), inference(avatar_component_clause, [], [f690])).
fof(f690, plain, (spl12_37 <=> (e10 = op1(e11, e12))), introduced(avatar_definition, [new_symbols(naming, [spl12_37])])).
fof(f1800, plain, (~ (h3(e10) = h3(op1(e11, e12))) | spl12_245), inference(avatar_component_clause, [], [f1798])).
fof(f1798, plain, (spl12_245 <=> (h3(e10) = h3(op1(e11, e12)))), introduced(avatar_definition, [new_symbols(naming, [spl12_245])])).
fof(f6248, plain, (~ spl12_24 | spl12_241), inference(avatar_contradiction_clause, [], [f6247])).
fof(f6247, plain, ($false | (~ spl12_24 | spl12_241)), inference(trivial_inequality_removal, [], [f6246])).
fof(f6246, plain, (~ (h3(e13) = h3(e13)) | (~ spl12_24 | spl12_241)), inference(forward_demodulation, [], [f1784, f636])).
fof(f636, plain, ((e13 = op1(e12, e12)) | ~ spl12_24), inference(avatar_component_clause, [], [f634])).
fof(f634, plain, (spl12_24 <=> (e13 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl12_24])])).
fof(f1784, plain, (~ (h3(e13) = h3(op1(e12, e12))) | spl12_241), inference(avatar_component_clause, [], [f1782])).
fof(f1782, plain, (spl12_241 <=> (h3(e13) = h3(op1(e12, e12)))), introduced(avatar_definition, [new_symbols(naming, [spl12_241])])).
fof(f6219, plain, (~ spl12_46 | ~ spl12_130 | ~ spl12_189 | spl12_247), inference(avatar_contradiction_clause, [], [f6218])).
fof(f6218, plain, ($false | (~ spl12_46 | ~ spl12_130 | ~ spl12_189 | spl12_247)), inference(subsumption_resolution, [], [f6217, f1087])).
fof(f1087, plain, ((e21 = op2(e21, e20)) | ~ spl12_130), inference(avatar_component_clause, [], [f1085])).
fof(f1085, plain, (spl12_130 <=> (e21 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl12_130])])).
fof(f6217, plain, (~ (e21 = op2(e21, e20)) | (~ spl12_46 | ~ spl12_189 | spl12_247)), inference(forward_demodulation, [], [f6216, f1339])).
fof(f6216, plain, (~ (op2(e21, e20) = h3(e11)) | (~ spl12_46 | ~ spl12_189 | spl12_247)), inference(forward_demodulation, [], [f6215, f730])).
fof(f730, plain, ((e11 = op1(e11, e10)) | ~ spl12_46), inference(avatar_component_clause, [], [f728])).
fof(f728, plain, (spl12_46 <=> (e11 = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl12_46])])).
fof(f6215, plain, (~ (op2(e21, e20) = h3(op1(e11, e10))) | (~ spl12_189 | spl12_247)), inference(forward_demodulation, [], [f1808, f1479])).
fof(f1808, plain, (~ (h3(op1(e11, e10)) = op2(e21, h3(e10))) | spl12_247), inference(avatar_component_clause, [], [f1806])).
fof(f1806, plain, (spl12_247 <=> (h3(op1(e11, e10)) = op2(e21, h3(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl12_247])])).
fof(f6190, plain, (~ spl12_31 | ~ spl12_115 | ~ spl12_189 | spl12_243), inference(avatar_contradiction_clause, [], [f6189])).
fof(f6189, plain, ($false | (~ spl12_31 | ~ spl12_115 | ~ spl12_189 | spl12_243)), inference(subsumption_resolution, [], [f6188, f1023])).
fof(f1023, plain, ((e22 = op2(e22, e20)) | ~ spl12_115), inference(avatar_component_clause, [], [f1021])).
fof(f1021, plain, (spl12_115 <=> (e22 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl12_115])])).
fof(f6188, plain, (~ (e22 = op2(e22, e20)) | (~ spl12_31 | ~ spl12_189 | spl12_243)), inference(forward_demodulation, [], [f6187, f464])).
fof(f464, plain, (e22 = h3(e12)), inference(cnf_transformation, [], [f28])).
fof(f6187, plain, (~ (op2(e22, e20) = h3(e12)) | (~ spl12_31 | ~ spl12_189 | spl12_243)), inference(forward_demodulation, [], [f6186, f666])).
fof(f666, plain, ((e12 = op1(e12, e10)) | ~ spl12_31), inference(avatar_component_clause, [], [f664])).
fof(f664, plain, (spl12_31 <=> (e12 = op1(e12, e10))), introduced(avatar_definition, [new_symbols(naming, [spl12_31])])).
fof(f6186, plain, (~ (op2(e22, e20) = h3(op1(e12, e10))) | (~ spl12_189 | spl12_243)), inference(forward_demodulation, [], [f1792, f1479])).
fof(f1792, plain, (~ (h3(op1(e12, e10)) = op2(e22, h3(e10))) | spl12_243), inference(avatar_component_clause, [], [f1790])).
fof(f1790, plain, (spl12_243 <=> (h3(op1(e12, e10)) = op2(e22, h3(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl12_243])])).
fof(f6158, plain, (~ spl12_11 | ~ spl12_95 | ~ spl12_237 | spl12_239), inference(avatar_contradiction_clause, [], [f6157])).
fof(f6157, plain, ($false | (~ spl12_11 | ~ spl12_95 | ~ spl12_237 | spl12_239)), inference(subsumption_resolution, [], [f6156, f938])).
fof(f938, plain, ((e22 = op2(e23, e21)) | ~ spl12_95), inference(avatar_component_clause, [], [f936])).
fof(f936, plain, (spl12_95 <=> (e22 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl12_95])])).
fof(f6156, plain, (~ (e22 = op2(e23, e21)) | (~ spl12_11 | ~ spl12_237 | spl12_239)), inference(forward_demodulation, [], [f6155, f464])).
fof(f6155, plain, (~ (op2(e23, e21) = h3(e12)) | (~ spl12_11 | ~ spl12_237 | spl12_239)), inference(forward_demodulation, [], [f6154, f581])).
fof(f581, plain, ((e12 = op1(e13, e11)) | ~ spl12_11), inference(avatar_component_clause, [], [f579])).
fof(f579, plain, (spl12_11 <=> (e12 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl12_11])])).
fof(f6154, plain, (~ (op2(e23, e21) = h3(op1(e13, e11))) | (~ spl12_237 | spl12_239)), inference(forward_demodulation, [], [f1776, f1767])).
fof(f1767, plain, ((e23 = h3(e13)) | ~ spl12_237), inference(avatar_component_clause, [], [f1766])).
fof(f1766, plain, (spl12_237 <=> (e23 = h3(e13))), introduced(avatar_definition, [new_symbols(naming, [spl12_237])])).
fof(f1776, plain, (~ (h3(op1(e13, e11)) = op2(h3(e13), e21)) | spl12_239), inference(avatar_component_clause, [], [f1774])).
fof(f1774, plain, (spl12_239 <=> (h3(op1(e13, e11)) = op2(h3(e13), e21))), introduced(avatar_definition, [new_symbols(naming, [spl12_239])])).
fof(f6129, plain, (~ spl12_1 | ~ spl12_85 | ~ spl12_189 | spl12_236 | ~ spl12_237), inference(avatar_contradiction_clause, [], [f6128])).
fof(f6128, plain, ($false | (~ spl12_1 | ~ spl12_85 | ~ spl12_189 | spl12_236 | ~ spl12_237)), inference(subsumption_resolution, [], [f6127, f896])).
fof(f896, plain, ((e20 = op2(e23, e23)) | ~ spl12_85), inference(avatar_component_clause, [], [f894])).
fof(f894, plain, (spl12_85 <=> (e20 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl12_85])])).
fof(f6127, plain, (~ (e20 = op2(e23, e23)) | (~ spl12_1 | ~ spl12_189 | spl12_236 | ~ spl12_237)), inference(forward_demodulation, [], [f6126, f1479])).
fof(f6126, plain, (~ (op2(e23, e23) = h3(e10)) | (~ spl12_1 | spl12_236 | ~ spl12_237)), inference(forward_demodulation, [], [f6125, f539])).
fof(f539, plain, ((e10 = op1(e13, e13)) | ~ spl12_1), inference(avatar_component_clause, [], [f537])).
fof(f537, plain, (spl12_1 <=> (e10 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl12_1])])).
fof(f6125, plain, (~ (op2(e23, e23) = h3(op1(e13, e13))) | (spl12_236 | ~ spl12_237)), inference(forward_demodulation, [], [f1764, f1767])).
fof(f1764, plain, (~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | spl12_236), inference(avatar_component_clause, [], [f1762])).
fof(f1762, plain, (spl12_236 <=> (h3(op1(e13, e13)) = op2(h3(e13), h3(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl12_236])])).
fof(f6097, plain, (~ spl12_55 | ~ spl12_119 | ~ spl12_237 | spl12_248), inference(avatar_contradiction_clause, [], [f6096])).
fof(f6096, plain, ($false | (~ spl12_55 | ~ spl12_119 | ~ spl12_237 | spl12_248)), inference(subsumption_resolution, [], [f6095, f1040])).
fof(f1040, plain, ((e22 = op2(e21, e23)) | ~ spl12_119), inference(avatar_component_clause, [], [f1038])).
fof(f1038, plain, (spl12_119 <=> (e22 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl12_119])])).
fof(f6095, plain, (~ (e22 = op2(e21, e23)) | (~ spl12_55 | ~ spl12_237 | spl12_248)), inference(forward_demodulation, [], [f6094, f464])).
fof(f6094, plain, (~ (op2(e21, e23) = h3(e12)) | (~ spl12_55 | ~ spl12_237 | spl12_248)), inference(forward_demodulation, [], [f6093, f768])).
fof(f768, plain, ((e12 = op1(e10, e12)) | ~ spl12_55), inference(avatar_component_clause, [], [f766])).
fof(f766, plain, (spl12_55 <=> (e12 = op1(e10, e12))), introduced(avatar_definition, [new_symbols(naming, [spl12_55])])).
fof(f6093, plain, (~ (op2(e21, e23) = h3(op1(e10, e12))) | (~ spl12_237 | spl12_248)), inference(forward_demodulation, [], [f1812, f1767])).
fof(f1812, plain, (~ (h3(op1(e10, e12)) = op2(e21, h3(e13))) | spl12_248), inference(avatar_component_clause, [], [f1810])).
fof(f1810, plain, (spl12_248 <=> (h3(op1(e10, e12)) = op2(e21, h3(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl12_248])])).
fof(f6063, plain, (~ spl12_44 | ~ spl12_237 | spl12_246 | ~ spl12_260), inference(avatar_contradiction_clause, [], [f6062])).
fof(f6062, plain, ($false | (~ spl12_44 | ~ spl12_237 | spl12_246 | ~ spl12_260)), inference(subsumption_resolution, [], [f6061, f1894])).
fof(f1894, plain, ((e23 = h2(e13)) | ~ spl12_260), inference(avatar_component_clause, [], [f1893])).
fof(f1893, plain, (spl12_260 <=> (e23 = h2(e13))), introduced(avatar_definition, [new_symbols(naming, [spl12_260])])).
fof(f6061, plain, (~ (e23 = h2(e13)) | (~ spl12_44 | ~ spl12_237 | spl12_246)), inference(forward_demodulation, [], [f6060, f1767])).
fof(f6060, plain, (~ (h2(e13) = h3(e13)) | (~ spl12_44 | spl12_246)), inference(forward_demodulation, [], [f1804, f721])).
fof(f721, plain, ((e13 = op1(e11, e11)) | ~ spl12_44), inference(avatar_component_clause, [], [f719])).
fof(f719, plain, (spl12_44 <=> (e13 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl12_44])])).
fof(f1804, plain, (~ (h2(e13) = h3(op1(e11, e11))) | spl12_246), inference(avatar_component_clause, [], [f1802])).
fof(f1802, plain, (spl12_246 <=> (h2(e13) = h3(op1(e11, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl12_246])])).
fof(f6030, plain, (~ spl12_35 | ~ spl12_119 | ~ spl12_237 | spl12_244), inference(avatar_contradiction_clause, [], [f6029])).
fof(f6029, plain, ($false | (~ spl12_35 | ~ spl12_119 | ~ spl12_237 | spl12_244)), inference(subsumption_resolution, [], [f6028, f1040])).
fof(f6028, plain, (~ (e22 = op2(e21, e23)) | (~ spl12_35 | ~ spl12_237 | spl12_244)), inference(forward_demodulation, [], [f6027, f464])).
fof(f6027, plain, (~ (op2(e21, e23) = h3(e12)) | (~ spl12_35 | ~ spl12_237 | spl12_244)), inference(forward_demodulation, [], [f6026, f683])).
fof(f683, plain, ((e12 = op1(e11, e13)) | ~ spl12_35), inference(avatar_component_clause, [], [f681])).
fof(f681, plain, (spl12_35 <=> (e12 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl12_35])])).
fof(f6026, plain, (~ (op2(e21, e23) = h3(op1(e11, e13))) | (~ spl12_237 | spl12_244)), inference(forward_demodulation, [], [f1796, f1767])).
fof(f1796, plain, (~ (h3(op1(e11, e13)) = op2(e21, h3(e13))) | spl12_244), inference(avatar_component_clause, [], [f1794])).
fof(f1794, plain, (spl12_244 <=> (h3(op1(e11, e13)) = op2(e21, h3(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl12_244])])).
fof(f5998, plain, (~ spl12_25 | ~ spl12_109 | ~ spl12_189 | spl12_242), inference(avatar_contradiction_clause, [], [f5997])).
fof(f5997, plain, ($false | (~ spl12_25 | ~ spl12_109 | ~ spl12_189 | spl12_242)), inference(subsumption_resolution, [], [f5996, f998])).
fof(f998, plain, ((e20 = op2(e22, e21)) | ~ spl12_109), inference(avatar_component_clause, [], [f996])).
fof(f996, plain, (spl12_109 <=> (e20 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl12_109])])).
fof(f5996, plain, (~ (e20 = op2(e22, e21)) | (~ spl12_25 | ~ spl12_189 | spl12_242)), inference(forward_demodulation, [], [f5995, f1479])).
fof(f5995, plain, (~ (op2(e22, e21) = h3(e10)) | (~ spl12_25 | spl12_242)), inference(forward_demodulation, [], [f1788, f641])).
fof(f641, plain, ((e10 = op1(e12, e11)) | ~ spl12_25), inference(avatar_component_clause, [], [f639])).
fof(f639, plain, (spl12_25 <=> (e10 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl12_25])])).
fof(f1788, plain, (~ (op2(e22, e21) = h3(op1(e12, e11))) | spl12_242), inference(avatar_component_clause, [], [f1786])).
fof(f1786, plain, (spl12_242 <=> (op2(e22, e21) = h3(op1(e12, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl12_242])])).
fof(f5968, plain, (~ spl12_18 | spl12_240), inference(avatar_contradiction_clause, [], [f5967])).
fof(f5967, plain, ($false | (~ spl12_18 | spl12_240)), inference(subsumption_resolution, [], [f5966, f1339])).
fof(f5966, plain, (~ (e21 = h3(e11)) | (~ spl12_18 | spl12_240)), inference(forward_demodulation, [], [f1780, f611])).
fof(f611, plain, ((e11 = op1(e12, e13)) | ~ spl12_18), inference(avatar_component_clause, [], [f609])).
fof(f609, plain, (spl12_18 <=> (e11 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl12_18])])).
fof(f1780, plain, (~ (e21 = h3(op1(e12, e13))) | spl12_240), inference(avatar_component_clause, [], [f1778])).
fof(f1778, plain, (spl12_240 <=> (e21 = h3(op1(e12, e13)))), introduced(avatar_definition, [new_symbols(naming, [spl12_240])])).
fof(f5934, plain, (~ spl12_6 | spl12_238), inference(avatar_contradiction_clause, [], [f5933])).
fof(f5933, plain, ($false | (~ spl12_6 | spl12_238)), inference(subsumption_resolution, [], [f5932, f1339])).
fof(f5932, plain, (~ (e21 = h3(e11)) | (~ spl12_6 | spl12_238)), inference(forward_demodulation, [], [f1772, f560])).
fof(f560, plain, ((e11 = op1(e13, e12)) | ~ spl12_6), inference(avatar_component_clause, [], [f558])).
fof(f558, plain, (spl12_6 <=> (e11 = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl12_6])])).
fof(f1772, plain, (~ (e21 = h3(op1(e13, e12))) | spl12_238), inference(avatar_component_clause, [], [f1770])).
fof(f1770, plain, (spl12_238 <=> (e21 = h3(op1(e13, e12)))), introduced(avatar_definition, [new_symbols(naming, [spl12_238])])).
fof(f5899, plain, (~ spl12_16 | ~ spl12_100 | ~ spl12_189 | spl12_235 | ~ spl12_237), inference(avatar_contradiction_clause, [], [f5898])).
fof(f5898, plain, ($false | (~ spl12_16 | ~ spl12_100 | ~ spl12_189 | spl12_235 | ~ spl12_237)), inference(subsumption_resolution, [], [f5897, f959])).
fof(f959, plain, ((e23 = op2(e23, e20)) | ~ spl12_100), inference(avatar_component_clause, [], [f957])).
fof(f957, plain, (spl12_100 <=> (e23 = op2(e23, e20))), introduced(avatar_definition, [new_symbols(naming, [spl12_100])])).
fof(f5897, plain, (~ (e23 = op2(e23, e20)) | (~ spl12_16 | ~ spl12_189 | spl12_235 | ~ spl12_237)), inference(forward_demodulation, [], [f5896, f1767])).
fof(f5896, plain, (~ (op2(e23, e20) = h3(e13)) | (~ spl12_16 | ~ spl12_189 | spl12_235 | ~ spl12_237)), inference(forward_demodulation, [], [f5895, f602])).
fof(f602, plain, ((e13 = op1(e13, e10)) | ~ spl12_16), inference(avatar_component_clause, [], [f600])).
fof(f600, plain, (spl12_16 <=> (e13 = op1(e13, e10))), introduced(avatar_definition, [new_symbols(naming, [spl12_16])])).
fof(f5895, plain, (~ (op2(e23, e20) = h3(op1(e13, e10))) | (~ spl12_189 | spl12_235 | ~ spl12_237)), inference(forward_demodulation, [], [f5894, f1767])).
fof(f5894, plain, (~ (h3(op1(e13, e10)) = op2(h3(e13), e20)) | (~ spl12_189 | spl12_235)), inference(forward_demodulation, [], [f1760, f1479])).
fof(f1760, plain, (~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | spl12_235), inference(avatar_component_clause, [], [f1758])).
fof(f1758, plain, (spl12_235 <=> (h3(op1(e13, e10)) = op2(h3(e13), h3(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl12_235])])).
fof(f5854, plain, (~ spl12_52 | ~ spl12_136 | ~ spl12_189 | spl12_234 | ~ spl12_237), inference(avatar_contradiction_clause, [], [f5853])).
fof(f5853, plain, ($false | (~ spl12_52 | ~ spl12_136 | ~ spl12_189 | spl12_234 | ~ spl12_237)), inference(subsumption_resolution, [], [f5852, f1112])).
fof(f1112, plain, ((e23 = op2(e20, e23)) | ~ spl12_136), inference(avatar_component_clause, [], [f1110])).
fof(f1110, plain, (spl12_136 <=> (e23 = op2(e20, e23))), introduced(avatar_definition, [new_symbols(naming, [spl12_136])])).
fof(f5852, plain, (~ (e23 = op2(e20, e23)) | (~ spl12_52 | ~ spl12_189 | spl12_234 | ~ spl12_237)), inference(forward_demodulation, [], [f5851, f1767])).
fof(f5851, plain, (~ (op2(e20, e23) = h3(e13)) | (~ spl12_52 | ~ spl12_189 | spl12_234 | ~ spl12_237)), inference(forward_demodulation, [], [f5850, f755])).
fof(f755, plain, ((e13 = op1(e10, e13)) | ~ spl12_52), inference(avatar_component_clause, [], [f753])).
fof(f753, plain, (spl12_52 <=> (e13 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl12_52])])).
fof(f5850, plain, (~ (op2(e20, e23) = h3(op1(e10, e13))) | (~ spl12_189 | spl12_234 | ~ spl12_237)), inference(forward_demodulation, [], [f5849, f1479])).
fof(f5849, plain, (~ (h3(op1(e10, e13)) = op2(h3(e10), e23)) | (spl12_234 | ~ spl12_237)), inference(forward_demodulation, [], [f1756, f1767])).
fof(f1756, plain, (~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | spl12_234), inference(avatar_component_clause, [], [f1754])).
fof(f1754, plain, (spl12_234 <=> (h3(op1(e10, e13)) = op2(h3(e10), h3(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl12_234])])).
fof(f5771, plain, (spl12_85 | ~ spl12_100 | ~ spl12_145 | ~ spl12_178), inference(avatar_split_clause, [], [f5765, f1417, f1149, f957, f894])).
fof(f1149, plain, (spl12_145 <=> (e20 = op2(e20, e20))), introduced(avatar_definition, [new_symbols(naming, [spl12_145])])).
fof(f1417, plain, (spl12_178 <=> (e20 = h4(e13))), introduced(avatar_definition, [new_symbols(naming, [spl12_178])])).
fof(f5765, plain, ((e20 = op2(e23, e23)) | (~ spl12_100 | ~ spl12_145 | ~ spl12_178)), inference(backward_demodulation, [], [f5644, f959])).
fof(f5644, plain, ((e20 = op2(e23, op2(e23, e20))) | (~ spl12_145 | ~ spl12_178)), inference(forward_demodulation, [], [f5630, f1151])).
fof(f1151, plain, ((e20 = op2(e20, e20)) | ~ spl12_145), inference(avatar_component_clause, [], [f1149])).
fof(f5630, plain, ((op2(e20, e20) = op2(e23, op2(e23, e20))) | ~ spl12_178), inference(backward_demodulation, [], [f1358, f1418])).
fof(f1418, plain, ((e20 = h4(e13)) | ~ spl12_178), inference(avatar_component_clause, [], [f1417])).
fof(f1358, plain, (op2(e23, op2(e23, e20)) = op2(h4(e13), e20)), inference(backward_demodulation, [], [f239, f471])).
fof(f471, plain, (op2(e23, e23) = h4(e13)), inference(cnf_transformation, [], [f29])).
fof(f29, plain, ((op2(e23, e23) = h4(e13)) & (op2(e23, op2(e23, e23)) = h4(e11)) & (h4(e10) = op2(op2(e23, op2(e23, e23)), e23)) & (e23 = h4(e12))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG016+1.p', ax29)).
fof(f239, plain, (op2(op2(e23, e23), e20) = op2(e23, op2(e23, e20))), inference(cnf_transformation, [], [f6])).
fof(f6, plain, ((op2(op2(e23, e23), e23) = op2(e23, op2(e23, e23))) & (op2(op2(e23, e23), e22) = op2(e23, op2(e23, e22))) & (op2(op2(e23, e23), e21) = op2(e23, op2(e23, e21))) & (op2(op2(e23, e23), e20) = op2(e23, op2(e23, e20))) & (op2(op2(e23, e22), e23) = op2(e23, op2(e22, e23))) & (op2(op2(e23, e22), e22) = op2(e23, op2(e22, e22))) & (op2(op2(e23, e22), e21) = op2(e23, op2(e22, e21))) & (op2(op2(e23, e22), e20) = op2(e23, op2(e22, e20))) & (op2(op2(e23, e21), e23) = op2(e23, op2(e21, e23))) & (op2(op2(e23, e21), e22) = op2(e23, op2(e21, e22))) & (op2(op2(e23, e21), e21) = op2(e23, op2(e21, e21))) & (op2(op2(e23, e21), e20) = op2(e23, op2(e21, e20))) & (op2(op2(e23, e20), e23) = op2(e23, op2(e20, e23))) & (op2(op2(e23, e20), e22) = op2(e23, op2(e20, e22))) & (op2(op2(e23, e20), e21) = op2(e23, op2(e20, e21))) & (op2(op2(e23, e20), e20) = op2(e23, op2(e20, e20))) & (op2(op2(e22, e23), e23) = op2(e22, op2(e23, e23))) & (op2(op2(e22, e23), e22) = op2(e22, op2(e23, e22))) & (op2(op2(e22, e23), e21) = op2(e22, op2(e23, e21))) & (op2(op2(e22, e23), e20) = op2(e22, op2(e23, e20))) & (op2(op2(e22, e22), e23) = op2(e22, op2(e22, e23))) & (op2(op2(e22, e22), e22) = op2(e22, op2(e22, e22))) & (op2(op2(e22, e22), e21) = op2(e22, op2(e22, e21))) & (op2(op2(e22, e22), e20) = op2(e22, op2(e22, e20))) & (op2(op2(e22, e21), e23) = op2(e22, op2(e21, e23))) & (op2(op2(e22, e21), e22) = op2(e22, op2(e21, e22))) & (op2(op2(e22, e21), e21) = op2(e22, op2(e21, e21))) & (op2(op2(e22, e21), e20) = op2(e22, op2(e21, e20))) & (op2(op2(e22, e20), e23) = op2(e22, op2(e20, e23))) & (op2(op2(e22, e20), e22) = op2(e22, op2(e20, e22))) & (op2(op2(e22, e20), e21) = op2(e22, op2(e20, e21))) & (op2(op2(e22, e20), e20) = op2(e22, op2(e20, e20))) & (op2(op2(e21, e23), e23) = op2(e21, op2(e23, e23))) & (op2(op2(e21, e23), e22) = op2(e21, op2(e23, e22))) & (op2(op2(e21, e23), e21) = op2(e21, op2(e23, e21))) & (op2(op2(e21, e23), e20) = op2(e21, op2(e23, e20))) & (op2(op2(e21, e22), e23) = op2(e21, op2(e22, e23))) & (op2(op2(e21, e22), e22) = op2(e21, op2(e22, e22))) & (op2(op2(e21, e22), e21) = op2(e21, op2(e22, e21))) & (op2(op2(e21, e22), e20) = op2(e21, op2(e22, e20))) & (op2(op2(e21, e21), e23) = op2(e21, op2(e21, e23))) & (op2(op2(e21, e21), e22) = op2(e21, op2(e21, e22))) & (op2(op2(e21, e21), e21) = op2(e21, op2(e21, e21))) & (op2(op2(e21, e21), e20) = op2(e21, op2(e21, e20))) & (op2(op2(e21, e20), e23) = op2(e21, op2(e20, e23))) & (op2(op2(e21, e20), e22) = op2(e21, op2(e20, e22))) & (op2(op2(e21, e20), e21) = op2(e21, op2(e20, e21))) & (op2(op2(e21, e20), e20) = op2(e21, op2(e20, e20))) & (op2(op2(e20, e23), e23) = op2(e20, op2(e23, e23))) & (op2(op2(e20, e23), e22) = op2(e20, op2(e23, e22))) & (op2(op2(e20, e23), e21) = op2(e20, op2(e23, e21))) & (op2(op2(e20, e23), e20) = op2(e20, op2(e23, e20))) & (op2(op2(e20, e22), e23) = op2(e20, op2(e22, e23))) & (op2(op2(e20, e22), e22) = op2(e20, op2(e22, e22))) & (op2(op2(e20, e22), e21) = op2(e20, op2(e22, e21))) & (op2(op2(e20, e22), e20) = op2(e20, op2(e22, e20))) & (op2(op2(e20, e21), e23) = op2(e20, op2(e21, e23))) & (op2(op2(e20, e21), e22) = op2(e20, op2(e21, e22))) & (op2(op2(e20, e21), e21) = op2(e20, op2(e21, e21))) & (op2(op2(e20, e21), e20) = op2(e20, op2(e21, e20))) & (op2(op2(e20, e20), e23) = op2(e20, op2(e20, e23))) & (op2(op2(e20, e20), e22) = op2(e20, op2(e20, e22))) & (op2(op2(e20, e20), e21) = op2(e20, op2(e20, e21))) & (op2(op2(e20, e20), e20) = op2(e20, op2(e20, e20)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG016+1.p', ax6)).
fof(f5729, plain, (spl12_100 | ~ spl12_149), inference(avatar_split_clause, [], [f5709, f1166, f957])).
fof(f1166, plain, (spl12_149 <=> (e20 = unit2)), introduced(avatar_definition, [new_symbols(naming, [spl12_149])])).
fof(f5709, plain, ((e23 = op2(e23, e20)) | ~ spl12_149), inference(backward_demodulation, [], [f250, f1168])).
fof(f1168, plain, ((e20 = unit2) | ~ spl12_149), inference(avatar_component_clause, [], [f1166])).
fof(f250, plain, (e23 = op2(e23, unit2)), inference(cnf_transformation, [], [f7])).
fof(f7, plain, (((e23 = unit2) | (e22 = unit2) | (e21 = unit2) | (e20 = unit2)) & (e23 = op2(e23, unit2)) & (e23 = op2(unit2, e23)) & (e22 = op2(e22, unit2)) & (e22 = op2(unit2, e22)) & (e21 = op2(e21, unit2)) & (e21 = op2(unit2, e21)) & (e20 = op2(e20, unit2)) & (e20 = op2(unit2, e20))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG016+1.p', ax7)).
fof(f5728, plain, (spl12_115 | ~ spl12_149), inference(avatar_split_clause, [], [f5708, f1166, f1021])).
fof(f5708, plain, ((e22 = op2(e22, e20)) | ~ spl12_149), inference(backward_demodulation, [], [f248, f1168])).
fof(f248, plain, (e22 = op2(e22, unit2)), inference(cnf_transformation, [], [f7])).
fof(f5727, plain, (spl12_142 | ~ spl12_149), inference(avatar_split_clause, [], [f5706, f1166, f1136])).
fof(f5706, plain, ((e21 = op2(e20, e21)) | ~ spl12_149), inference(backward_demodulation, [], [f245, f1168])).
fof(f245, plain, (e21 = op2(unit2, e21)), inference(cnf_transformation, [], [f7])).
fof(f5436, plain, (spl12_44 | ~ spl12_6 | ~ spl12_16 | ~ spl12_24 | ~ spl12_37 | ~ spl12_46 | ~ spl12_52), inference(avatar_split_clause, [], [f5435, f753, f728, f690, f634, f600, f558, f719])).
fof(f5435, plain, ((e13 = op1(e11, e11)) | (~ spl12_6 | ~ spl12_16 | ~ spl12_24 | ~ spl12_37 | ~ spl12_46 | ~ spl12_52)), inference(forward_demodulation, [], [f5434, f602])).
fof(f5434, plain, ((op1(e11, e11) = op1(e13, e10)) | (~ spl12_6 | ~ spl12_24 | ~ spl12_37 | ~ spl12_46 | ~ spl12_52)), inference(forward_demodulation, [], [f5433, f755])).
fof(f5433, plain, ((op1(e11, e11) = op1(op1(e10, e13), e10)) | (~ spl12_6 | ~ spl12_24 | ~ spl12_37 | ~ spl12_46)), inference(forward_demodulation, [], [f4944, f730])).
fof(f4944, plain, ((op1(op1(e10, e13), e10) = op1(e11, op1(e11, e10))) | (~ spl12_6 | ~ spl12_24 | ~ spl12_37)), inference(forward_demodulation, [], [f98, f3565])).
fof(f3565, plain, ((op1(e10, e13) = op1(e11, e11)) | (~ spl12_6 | ~ spl12_24 | ~ spl12_37)), inference(forward_demodulation, [], [f3564, f3024])).
fof(f3024, plain, ((op1(e10, e13) = op1(op1(e10, e12), e12)) | ~ spl12_24), inference(forward_demodulation, [], [f88, f636])).
fof(f88, plain, (op1(op1(e10, e12), e12) = op1(e10, op1(e12, e12))), inference(cnf_transformation, [], [f2])).
fof(f2, plain, ((op1(op1(e13, e13), e13) = op1(e13, op1(e13, e13))) & (op1(op1(e13, e13), e12) = op1(e13, op1(e13, e12))) & (op1(op1(e13, e13), e11) = op1(e13, op1(e13, e11))) & (op1(op1(e13, e13), e10) = op1(e13, op1(e13, e10))) & (op1(op1(e13, e12), e13) = op1(e13, op1(e12, e13))) & (op1(op1(e13, e12), e12) = op1(e13, op1(e12, e12))) & (op1(op1(e13, e12), e11) = op1(e13, op1(e12, e11))) & (op1(op1(e13, e12), e10) = op1(e13, op1(e12, e10))) & (op1(op1(e13, e11), e13) = op1(e13, op1(e11, e13))) & (op1(op1(e13, e11), e12) = op1(e13, op1(e11, e12))) & (op1(op1(e13, e11), e11) = op1(e13, op1(e11, e11))) & (op1(op1(e13, e11), e10) = op1(e13, op1(e11, e10))) & (op1(op1(e13, e10), e13) = op1(e13, op1(e10, e13))) & (op1(op1(e13, e10), e12) = op1(e13, op1(e10, e12))) & (op1(op1(e13, e10), e11) = op1(e13, op1(e10, e11))) & (op1(op1(e13, e10), e10) = op1(e13, op1(e10, e10))) & (op1(op1(e12, e13), e13) = op1(e12, op1(e13, e13))) & (op1(op1(e12, e13), e12) = op1(e12, op1(e13, e12))) & (op1(op1(e12, e13), e11) = op1(e12, op1(e13, e11))) & (op1(op1(e12, e13), e10) = op1(e12, op1(e13, e10))) & (op1(op1(e12, e12), e13) = op1(e12, op1(e12, e13))) & (op1(op1(e12, e12), e12) = op1(e12, op1(e12, e12))) & (op1(op1(e12, e12), e11) = op1(e12, op1(e12, e11))) & (op1(op1(e12, e12), e10) = op1(e12, op1(e12, e10))) & (op1(op1(e12, e11), e13) = op1(e12, op1(e11, e13))) & (op1(op1(e12, e11), e12) = op1(e12, op1(e11, e12))) & (op1(op1(e12, e11), e11) = op1(e12, op1(e11, e11))) & (op1(op1(e12, e11), e10) = op1(e12, op1(e11, e10))) & (op1(op1(e12, e10), e13) = op1(e12, op1(e10, e13))) & (op1(op1(e12, e10), e12) = op1(e12, op1(e10, e12))) & (op1(op1(e12, e10), e11) = op1(e12, op1(e10, e11))) & (op1(op1(e12, e10), e10) = op1(e12, op1(e10, e10))) & (op1(op1(e11, e13), e13) = op1(e11, op1(e13, e13))) & (op1(op1(e11, e13), e12) = op1(e11, op1(e13, e12))) & (op1(op1(e11, e13), e11) = op1(e11, op1(e13, e11))) & (op1(op1(e11, e13), e10) = op1(e11, op1(e13, e10))) & (op1(op1(e11, e12), e13) = op1(e11, op1(e12, e13))) & (op1(op1(e11, e12), e12) = op1(e11, op1(e12, e12))) & (op1(op1(e11, e12), e11) = op1(e11, op1(e12, e11))) & (op1(op1(e11, e12), e10) = op1(e11, op1(e12, e10))) & (op1(op1(e11, e11), e13) = op1(e11, op1(e11, e13))) & (op1(op1(e11, e11), e12) = op1(e11, op1(e11, e12))) & (op1(op1(e11, e11), e11) = op1(e11, op1(e11, e11))) & (op1(op1(e11, e11), e10) = op1(e11, op1(e11, e10))) & (op1(op1(e11, e10), e13) = op1(e11, op1(e10, e13))) & (op1(op1(e11, e10), e12) = op1(e11, op1(e10, e12))) & (op1(op1(e11, e10), e11) = op1(e11, op1(e10, e11))) & (op1(op1(e11, e10), e10) = op1(e11, op1(e10, e10))) & (op1(op1(e10, e13), e13) = op1(e10, op1(e13, e13))) & (op1(op1(e10, e13), e12) = op1(e10, op1(e13, e12))) & (op1(op1(e10, e13), e11) = op1(e10, op1(e13, e11))) & (op1(op1(e10, e13), e10) = op1(e10, op1(e13, e10))) & (op1(op1(e10, e12), e13) = op1(e10, op1(e12, e13))) & (op1(op1(e10, e12), e12) = op1(e10, op1(e12, e12))) & (op1(op1(e10, e12), e11) = op1(e10, op1(e12, e11))) & (op1(op1(e10, e12), e10) = op1(e10, op1(e12, e10))) & (op1(op1(e10, e11), e13) = op1(e10, op1(e11, e13))) & (op1(op1(e10, e11), e12) = op1(e10, op1(e11, e12))) & (op1(op1(e10, e11), e11) = op1(e10, op1(e11, e11))) & (op1(op1(e10, e11), e10) = op1(e10, op1(e11, e10))) & (op1(op1(e10, e10), e13) = op1(e10, op1(e10, e13))) & (op1(op1(e10, e10), e12) = op1(e10, op1(e10, e12))) & (op1(op1(e10, e10), e11) = op1(e10, op1(e10, e11))) & (op1(op1(e10, e10), e10) = op1(e10, op1(e10, e10)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG016+1.p', ax2)).
fof(f3564, plain, ((op1(e11, e11) = op1(op1(e10, e12), e12)) | (~ spl12_6 | ~ spl12_24 | ~ spl12_37)), inference(forward_demodulation, [], [f2987, f560])).
fof(f2987, plain, ((op1(op1(e10, e12), e12) = op1(e11, op1(e13, e12))) | (~ spl12_24 | ~ spl12_37)), inference(backward_demodulation, [], [f108, f2984])).
fof(f2984, plain, ((op1(e10, e12) = op1(e11, e13)) | (~ spl12_24 | ~ spl12_37)), inference(forward_demodulation, [], [f2983, f692])).
fof(f2983, plain, ((op1(e11, e13) = op1(op1(e11, e12), e12)) | ~ spl12_24), inference(forward_demodulation, [], [f104, f636])).
fof(f104, plain, (op1(op1(e11, e12), e12) = op1(e11, op1(e12, e12))), inference(cnf_transformation, [], [f2])).
fof(f108, plain, (op1(op1(e11, e13), e12) = op1(e11, op1(e13, e12))), inference(cnf_transformation, [], [f2])).
fof(f98, plain, (op1(op1(e11, e11), e10) = op1(e11, op1(e11, e10))), inference(cnf_transformation, [], [f2])).
fof(f5418, plain, (spl12_81 | ~ spl12_65), inference(avatar_split_clause, [], [f5360, f809, f877])).
fof(f877, plain, (spl12_81 <=> (e10 = inv1(e10))), introduced(avatar_definition, [new_symbols(naming, [spl12_81])])).
fof(f809, plain, (spl12_65 <=> (e10 = unit1)), introduced(avatar_definition, [new_symbols(naming, [spl12_65])])).
fof(f5360, plain, ((e10 = inv1(e10)) | ~ spl12_65), inference(backward_demodulation, [], [f264, f811])).
fof(f811, plain, ((e10 = unit1) | ~ spl12_65), inference(avatar_component_clause, [], [f809])).
fof(f264, plain, (unit1 = inv1(unit1)), inference(cnf_transformation, [], [f9])).
fof(f9, plain, (unit1 = inv1(unit1)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG016+1.p', ax9)).
fof(f5383, plain, (~ spl12_65 | ~ spl12_81 | ~ spl12_189 | ~ spl12_211 | spl12_233), inference(avatar_contradiction_clause, [], [f5382])).
fof(f5382, plain, ($false | (~ spl12_65 | ~ spl12_81 | ~ spl12_189 | ~ spl12_211 | spl12_233)), inference(subsumption_resolution, [], [f5372, f1479])).
fof(f5372, plain, (~ (e20 = h3(e10)) | (~ spl12_65 | ~ spl12_81 | ~ spl12_189 | ~ spl12_211 | spl12_233)), inference(backward_demodulation, [], [f5208, f811])).
fof(f5208, plain, (~ (e20 = h3(unit1)) | (~ spl12_81 | ~ spl12_189 | ~ spl12_211 | spl12_233)), inference(backward_demodulation, [], [f3878, f1589])).
fof(f1589, plain, ((e20 = h1(e13)) | ~ spl12_211), inference(avatar_component_clause, [], [f1588])).
fof(f1588, plain, (spl12_211 <=> (e20 = h1(e13))), introduced(avatar_definition, [new_symbols(naming, [spl12_211])])).
fof(f3878, plain, (~ (h1(e13) = h3(unit1)) | (~ spl12_81 | ~ spl12_189 | spl12_233)), inference(backward_demodulation, [], [f3087, f459])).
fof(f459, plain, (op2(e20, e20) = h1(e13)), inference(cnf_transformation, [], [f26])).
fof(f26, plain, ((op2(e20, e20) = h1(e13)) & (op2(e20, op2(e20, e20)) = h1(e11)) & (h1(e10) = op2(op2(e20, op2(e20, e20)), e20)) & (e20 = h1(e12))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG016+1.p', ax26)).
fof(f3087, plain, (~ (op2(e20, e20) = h3(unit1)) | (~ spl12_81 | ~ spl12_189 | spl12_233)), inference(backward_demodulation, [], [f2500, f1479])).
fof(f2500, plain, (~ (op2(h3(e10), h3(e10)) = h3(unit1)) | (~ spl12_81 | spl12_233)), inference(backward_demodulation, [], [f1752, f2478])).
fof(f2478, plain, ((op1(e10, e10) = unit1) | ~ spl12_81), inference(backward_demodulation, [], [f151, f879])).
fof(f879, plain, ((e10 = inv1(e10)) | ~ spl12_81), inference(avatar_component_clause, [], [f877])).
fof(f151, plain, (unit1 = op1(e10, inv1(e10))), inference(cnf_transformation, [], [f4])).
fof(f4, plain, (((e13 = inv1(e13)) | (e12 = inv1(e13)) | (e11 = inv1(e13)) | (e10 = inv1(e13))) & ((e13 = inv1(e12)) | (e12 = inv1(e12)) | (e11 = inv1(e12)) | (e10 = inv1(e12))) & ((e13 = inv1(e11)) | (e12 = inv1(e11)) | (e11 = inv1(e11)) | (e10 = inv1(e11))) & ((e13 = inv1(e10)) | (e12 = inv1(e10)) | (e11 = inv1(e10)) | (e10 = inv1(e10))) & (unit1 = op1(inv1(e13), e13)) & (unit1 = op1(e13, inv1(e13))) & (unit1 = op1(inv1(e12), e12)) & (unit1 = op1(e12, inv1(e12))) & (unit1 = op1(inv1(e11), e11)) & (unit1 = op1(e11, inv1(e11))) & (unit1 = op1(inv1(e10), e10)) & (unit1 = op1(e10, inv1(e10)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG016+1.p', ax4)).
fof(f1752, plain, (~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10))) | spl12_233), inference(avatar_component_clause, [], [f1750])).
fof(f1750, plain, (spl12_233 <=> (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl12_233])])).
fof(f5378, plain, (spl12_16 | ~ spl12_65), inference(avatar_split_clause, [], [f5359, f809, f600])).
fof(f5359, plain, ((e13 = op1(e13, e10)) | ~ spl12_65), inference(backward_demodulation, [], [f149, f811])).
fof(f149, plain, (e13 = op1(e13, unit1)), inference(cnf_transformation, [], [f3])).
fof(f3, plain, (((e13 = unit1) | (e12 = unit1) | (e11 = unit1) | (e10 = unit1)) & (e13 = op1(e13, unit1)) & (e13 = op1(unit1, e13)) & (e12 = op1(e12, unit1)) & (e12 = op1(unit1, e12)) & (e11 = op1(e11, unit1)) & (e11 = op1(unit1, e11)) & (e10 = op1(e10, unit1)) & (e10 = op1(unit1, e10))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG016+1.p', ax3)).
fof(f5377, plain, (spl12_52 | ~ spl12_65), inference(avatar_split_clause, [], [f5358, f809, f753])).
fof(f5358, plain, ((e13 = op1(e10, e13)) | ~ spl12_65), inference(backward_demodulation, [], [f148, f811])).
fof(f148, plain, (e13 = op1(unit1, e13)), inference(cnf_transformation, [], [f3])).
fof(f5376, plain, (spl12_46 | ~ spl12_65), inference(avatar_split_clause, [], [f5355, f809, f728])).
fof(f5355, plain, ((e11 = op1(e11, e10)) | ~ spl12_65), inference(backward_demodulation, [], [f145, f811])).
fof(f145, plain, (e11 = op1(e11, unit1)), inference(cnf_transformation, [], [f3])).
fof(f5375, plain, (spl12_58 | ~ spl12_65), inference(avatar_split_clause, [], [f5354, f809, f779])).
fof(f5354, plain, ((e11 = op1(e10, e11)) | ~ spl12_65), inference(backward_demodulation, [], [f144, f811])).
fof(f144, plain, (e11 = op1(unit1, e11)), inference(cnf_transformation, [], [f3])).
fof(f5342, plain, (spl12_260 | ~ spl12_102 | ~ spl12_136 | ~ spl12_189), inference(avatar_split_clause, [], [f5341, f1478, f1110, f966, f1893])).
fof(f966, plain, (spl12_102 <=> (e21 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl12_102])])).
fof(f5341, plain, ((e23 = h2(e13)) | (~ spl12_102 | ~ spl12_136 | ~ spl12_189)), inference(backward_demodulation, [], [f5105, f1112])).
fof(f5105, plain, ((op2(e20, e23) = h2(e13)) | (~ spl12_102 | ~ spl12_189)), inference(forward_demodulation, [], [f5104, f463])).
fof(f463, plain, (op2(e21, e21) = h2(e13)), inference(cnf_transformation, [], [f27])).
fof(f27, plain, ((op2(e21, e21) = h2(e13)) & (op2(e21, op2(e21, e21)) = h2(e11)) & (h2(e10) = op2(op2(e21, op2(e21, e21)), e21)) & (e21 = h2(e12))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG016+1.p', ax27)).
fof(f5104, plain, ((op2(e20, e23) = op2(e21, e21)) | (~ spl12_102 | ~ spl12_189)), inference(forward_demodulation, [], [f3075, f968])).
fof(f968, plain, ((e21 = op2(e22, e23)) | ~ spl12_102), inference(avatar_component_clause, [], [f966])).
fof(f3075, plain, ((op2(e20, e23) = op2(e21, op2(e22, e23))) | ~ spl12_189), inference(backward_demodulation, [], [f1345, f1479])).
fof(f1345, plain, (op2(e21, op2(e22, e23)) = op2(h3(e10), e23)), inference(backward_demodulation, [], [f206, f1341])).
fof(f1341, plain, (op2(e21, e22) = h3(e10)), inference(forward_demodulation, [], [f1340, f1336])).
fof(f1340, plain, (h3(e10) = op2(op2(e22, h3(e13)), e22)), inference(forward_demodulation, [], [f465, f467])).
fof(f465, plain, (op2(op2(e22, op2(e22, e22)), e22) = h3(e10)), inference(cnf_transformation, [], [f28])).
fof(f206, plain, (op2(op2(e21, e22), e23) = op2(e21, op2(e22, e23))), inference(cnf_transformation, [], [f6])).
fof(f5314, plain, (~ spl12_90 | ~ spl12_152), inference(avatar_contradiction_clause, [], [f5313])).
fof(f5313, plain, ($false | (~ spl12_90 | ~ spl12_152)), inference(subsumption_resolution, [], [f5312, f423])).
fof(f423, plain, ~ (e21 = e22), inference(cnf_transformation, [], [f20])).
fof(f20, plain, (~ (e22 = e23) & ~ (e21 = e23) & ~ (e21 = e22) & ~ (e20 = e23) & ~ (e20 = e22) & ~ (e20 = e21)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG016+1.p', ax20)).
fof(f5312, plain, ((e21 = e22) | (~ spl12_90 | ~ spl12_152)), inference(forward_demodulation, [], [f5293, f917])).
fof(f917, plain, ((e21 = op2(e23, e22)) | ~ spl12_90), inference(avatar_component_clause, [], [f915])).
fof(f915, plain, (spl12_90 <=> (e21 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl12_90])])).
fof(f5293, plain, ((e22 = op2(e23, e22)) | ~ spl12_152), inference(backward_demodulation, [], [f247, f1180])).
fof(f1180, plain, ((e23 = unit2) | ~ spl12_152), inference(avatar_component_clause, [], [f1178])).
fof(f1178, plain, (spl12_152 <=> (e23 = unit2)), introduced(avatar_definition, [new_symbols(naming, [spl12_152])])).
fof(f247, plain, (e22 = op2(unit2, e22)), inference(cnf_transformation, [], [f7])).
fof(f5165, plain, (spl12_178 | ~ spl12_90 | ~ spl12_121 | ~ spl12_237), inference(avatar_split_clause, [], [f5164, f1766, f1047, f915, f1417])).
fof(f1047, plain, (spl12_121 <=> (e20 = op2(e21, e22))), introduced(avatar_definition, [new_symbols(naming, [spl12_121])])).
fof(f5164, plain, ((e20 = h4(e13)) | (~ spl12_90 | ~ spl12_121 | ~ spl12_237)), inference(forward_demodulation, [], [f5163, f1049])).
fof(f1049, plain, ((e20 = op2(e21, e22)) | ~ spl12_121), inference(avatar_component_clause, [], [f1047])).
fof(f5163, plain, ((op2(e21, e22) = h4(e13)) | (~ spl12_90 | ~ spl12_237)), inference(forward_demodulation, [], [f5162, f471])).
fof(f5162, plain, ((op2(e21, e22) = op2(e23, e23)) | (~ spl12_90 | ~ spl12_237)), inference(forward_demodulation, [], [f5161, f917])).
fof(f5161, plain, ((op2(e23, e23) = op2(op2(e23, e22), e22)) | ~ spl12_237), inference(forward_demodulation, [], [f1329, f1767])).
fof(f1329, plain, (op2(op2(e23, e22), e22) = op2(e23, h3(e13))), inference(backward_demodulation, [], [f237, f467])).
fof(f237, plain, (op2(op2(e23, e22), e22) = op2(e23, op2(e22, e22))), inference(cnf_transformation, [], [f6])).
fof(f4768, plain, (~ spl12_18 | ~ spl12_68), inference(avatar_contradiction_clause, [], [f4767])).
fof(f4767, plain, ($false | (~ spl12_18 | ~ spl12_68)), inference(subsumption_resolution, [], [f4766, f417])).
fof(f417, plain, ~ (e11 = e12), inference(cnf_transformation, [], [f19])).
fof(f19, plain, (~ (e12 = e13) & ~ (e11 = e13) & ~ (e11 = e12) & ~ (e10 = e13) & ~ (e10 = e12) & ~ (e10 = e11)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG016+1.p', ax19)).
fof(f4766, plain, ((e11 = e12) | (~ spl12_18 | ~ spl12_68)), inference(forward_demodulation, [], [f4744, f611])).
fof(f4744, plain, ((e12 = op1(e12, e13)) | ~ spl12_68), inference(backward_demodulation, [], [f147, f823])).
fof(f823, plain, ((e13 = unit1) | ~ spl12_68), inference(avatar_component_clause, [], [f821])).
fof(f821, plain, (spl12_68 <=> (e13 = unit1)), introduced(avatar_definition, [new_symbols(naming, [spl12_68])])).
fof(f147, plain, (e12 = op1(e12, unit1)), inference(cnf_transformation, [], [f3])).
fof(f4706, plain, (spl12_31 | ~ spl12_37 | ~ spl12_79), inference(avatar_split_clause, [], [f4705, f868, f690, f664])).
fof(f868, plain, (spl12_79 <=> (e12 = inv1(e11))), introduced(avatar_definition, [new_symbols(naming, [spl12_79])])).
fof(f4705, plain, ((e12 = op1(e12, e10)) | (~ spl12_37 | ~ spl12_79)), inference(forward_demodulation, [], [f4693, f146])).
fof(f146, plain, (e12 = op1(unit1, e12)), inference(cnf_transformation, [], [f3])).
fof(f4693, plain, ((op1(e12, e10) = op1(unit1, e12)) | (~ spl12_37 | ~ spl12_79)), inference(backward_demodulation, [], [f2732, f4681])).
fof(f4681, plain, ((op1(e12, e11) = unit1) | ~ spl12_79), inference(backward_demodulation, [], [f154, f870])).
fof(f870, plain, ((e12 = inv1(e11)) | ~ spl12_79), inference(avatar_component_clause, [], [f868])).
fof(f154, plain, (unit1 = op1(inv1(e11), e11)), inference(cnf_transformation, [], [f4])).
fof(f2732, plain, ((op1(e12, e10) = op1(op1(e12, e11), e12)) | ~ spl12_37), inference(backward_demodulation, [], [f116, f692])).
fof(f116, plain, (op1(op1(e12, e11), e12) = op1(e12, op1(e11, e12))), inference(cnf_transformation, [], [f2])).
fof(f4535, plain, (~ spl12_151 | spl12_159), inference(avatar_contradiction_clause, [], [f4534])).
fof(f4534, plain, ($false | (~ spl12_151 | spl12_159)), inference(subsumption_resolution, [], [f4505, f1209])).
fof(f1209, plain, (~ (e22 = inv2(e22)) | spl12_159), inference(avatar_component_clause, [], [f1208])).
fof(f1208, plain, (spl12_159 <=> (e22 = inv2(e22))), introduced(avatar_definition, [new_symbols(naming, [spl12_159])])).
fof(f4505, plain, ((e22 = inv2(e22)) | ~ spl12_151), inference(backward_demodulation, [], [f291, f1176])).
fof(f1176, plain, ((e22 = unit2) | ~ spl12_151), inference(avatar_component_clause, [], [f1174])).
fof(f1174, plain, (spl12_151 <=> (e22 = unit2)), introduced(avatar_definition, [new_symbols(naming, [spl12_151])])).
fof(f291, plain, (unit2 = inv2(unit2)), inference(cnf_transformation, [], [f13])).
fof(f13, plain, (unit2 = inv2(unit2)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG016+1.p', ax13)).
fof(f4212, plain, (spl12_1 | ~ spl12_6 | ~ spl12_24 | ~ spl12_37), inference(avatar_split_clause, [], [f3538, f690, f634, f558, f537])).
fof(f3538, plain, ((e10 = op1(e13, e13)) | (~ spl12_6 | ~ spl12_24 | ~ spl12_37)), inference(forward_demodulation, [], [f3537, f692])).
fof(f3537, plain, ((op1(e11, e12) = op1(e13, e13)) | (~ spl12_6 | ~ spl12_24)), inference(forward_demodulation, [], [f2952, f560])).
fof(f2952, plain, ((op1(e13, e13) = op1(op1(e13, e12), e12)) | ~ spl12_24), inference(forward_demodulation, [], [f136, f636])).
fof(f136, plain, (op1(op1(e13, e12), e12) = op1(e13, op1(e12, e12))), inference(cnf_transformation, [], [f2])).
fof(f4161, plain, (~ spl12_67 | spl12_75), inference(avatar_contradiction_clause, [], [f4160])).
fof(f4160, plain, ($false | (~ spl12_67 | spl12_75)), inference(subsumption_resolution, [], [f4138, f852])).
fof(f852, plain, (~ (e12 = inv1(e12)) | spl12_75), inference(avatar_component_clause, [], [f851])).
fof(f851, plain, (spl12_75 <=> (e12 = inv1(e12))), introduced(avatar_definition, [new_symbols(naming, [spl12_75])])).
fof(f4138, plain, ((e12 = inv1(e12)) | ~ spl12_67), inference(backward_demodulation, [], [f264, f819])).
fof(f819, plain, ((e12 = unit1) | ~ spl12_67), inference(avatar_component_clause, [], [f817])).
fof(f817, plain, (spl12_67 <=> (e12 = unit1)), introduced(avatar_definition, [new_symbols(naming, [spl12_67])])).
fof(f4017, plain, (~ spl12_121 | ~ spl12_150), inference(avatar_contradiction_clause, [], [f4016])).
fof(f4016, plain, ($false | (~ spl12_121 | ~ spl12_150)), inference(subsumption_resolution, [], [f4015, f421])).
fof(f421, plain, ~ (e20 = e22), inference(cnf_transformation, [], [f20])).
fof(f4015, plain, ((e20 = e22) | (~ spl12_121 | ~ spl12_150)), inference(forward_demodulation, [], [f3996, f1049])).
fof(f3996, plain, ((e22 = op2(e21, e22)) | ~ spl12_150), inference(backward_demodulation, [], [f247, f1172])).
fof(f1172, plain, ((e21 = unit2) | ~ spl12_150), inference(avatar_component_clause, [], [f1170])).
fof(f1170, plain, (spl12_150 <=> (e21 = unit2)), introduced(avatar_definition, [new_symbols(naming, [spl12_150])])).
fof(f3982, plain, (spl12_90 | ~ spl12_237), inference(avatar_split_clause, [], [f3973, f1766, f915])).
fof(f3973, plain, ((e21 = op2(e23, e22)) | ~ spl12_237), inference(backward_demodulation, [], [f1337, f1767])).
fof(f1337, plain, (e21 = op2(h3(e13), e22)), inference(backward_demodulation, [], [f1287, f467])).
fof(f1287, plain, (e21 = op2(op2(e22, e22), e22)), inference(backward_demodulation, [], [f221, f454])).
fof(f221, plain, (op2(op2(e22, e22), e22) = op2(e22, op2(e22, e22))), inference(cnf_transformation, [], [f6])).
fof(f3981, plain, (spl12_102 | ~ spl12_237), inference(avatar_split_clause, [], [f3972, f1766, f966])).
fof(f3972, plain, ((e21 = op2(e22, e23)) | ~ spl12_237), inference(backward_demodulation, [], [f1336, f1767])).
fof(f3757, plain, (spl12_66 | ~ spl12_18 | ~ spl12_76), inference(avatar_split_clause, [], [f3756, f855, f609, f813])).
fof(f813, plain, (spl12_66 <=> (e11 = unit1)), introduced(avatar_definition, [new_symbols(naming, [spl12_66])])).
fof(f855, plain, (spl12_76 <=> (e13 = inv1(e12))), introduced(avatar_definition, [new_symbols(naming, [spl12_76])])).
fof(f3756, plain, ((e11 = unit1) | (~ spl12_18 | ~ spl12_76)), inference(forward_demodulation, [], [f3342, f611])).
fof(f3342, plain, ((op1(e12, e13) = unit1) | ~ spl12_76), inference(backward_demodulation, [], [f155, f857])).
fof(f857, plain, ((e13 = inv1(e12)) | ~ spl12_76), inference(avatar_component_clause, [], [f855])).
fof(f155, plain, (unit1 = op1(e12, inv1(e12))), inference(cnf_transformation, [], [f4])).
fof(f3563, plain, (~ spl12_54 | ~ spl12_18 | ~ spl12_24 | ~ spl12_37), inference(avatar_split_clause, [], [f3562, f690, f634, f609, f762])).
fof(f762, plain, (spl12_54 <=> (e11 = op1(e10, e12))), introduced(avatar_definition, [new_symbols(naming, [spl12_54])])).
fof(f3562, plain, (~ (e11 = op1(e10, e12)) | (~ spl12_18 | ~ spl12_24 | ~ spl12_37)), inference(forward_demodulation, [], [f2989, f611])).
fof(f2989, plain, (~ (op1(e10, e12) = op1(e12, e13)) | (~ spl12_24 | ~ spl12_37)), inference(backward_demodulation, [], [f339, f2984])).
fof(f339, plain, ~ (op1(e11, e13) = op1(e12, e13)), inference(cnf_transformation, [], [f17])).
fof(f17, plain, (~ (op1(e13, e12) = op1(e13, e13)) & ~ (op1(e13, e11) = op1(e13, e13)) & ~ (op1(e13, e11) = op1(e13, e12)) & ~ (op1(e13, e10) = op1(e13, e13)) & ~ (op1(e13, e10) = op1(e13, e12)) & ~ (op1(e13, e10) = op1(e13, e11)) & ~ (op1(e12, e12) = op1(e12, e13)) & ~ (op1(e12, e11) = op1(e12, e13)) & ~ (op1(e12, e11) = op1(e12, e12)) & ~ (op1(e12, e10) = op1(e12, e13)) & ~ (op1(e12, e10) = op1(e12, e12)) & ~ (op1(e12, e10) = op1(e12, e11)) & ~ (op1(e11, e12) = op1(e11, e13)) & ~ (op1(e11, e11) = op1(e11, e13)) & ~ (op1(e11, e11) = op1(e11, e12)) & ~ (op1(e11, e10) = op1(e11, e13)) & ~ (op1(e11, e10) = op1(e11, e12)) & ~ (op1(e11, e10) = op1(e11, e11)) & ~ (op1(e10, e12) = op1(e10, e13)) & ~ (op1(e10, e11) = op1(e10, e13)) & ~ (op1(e10, e11) = op1(e10, e12)) & ~ (op1(e10, e10) = op1(e10, e13)) & ~ (op1(e10, e10) = op1(e10, e12)) & ~ (op1(e10, e10) = op1(e10, e11)) & ~ (op1(e12, e13) = op1(e13, e13)) & ~ (op1(e11, e13) = op1(e13, e13)) & ~ (op1(e11, e13) = op1(e12, e13)) & ~ (op1(e10, e13) = op1(e13, e13)) & ~ (op1(e10, e13) = op1(e12, e13)) & ~ (op1(e10, e13) = op1(e11, e13)) & ~ (op1(e12, e12) = op1(e13, e12)) & ~ (op1(e11, e12) = op1(e13, e12)) & ~ (op1(e11, e12) = op1(e12, e12)) & ~ (op1(e10, e12) = op1(e13, e12)) & ~ (op1(e10, e12) = op1(e12, e12)) & ~ (op1(e10, e12) = op1(e11, e12)) & ~ (op1(e12, e11) = op1(e13, e11)) & ~ (op1(e11, e11) = op1(e13, e11)) & ~ (op1(e11, e11) = op1(e12, e11)) & ~ (op1(e10, e11) = op1(e13, e11)) & ~ (op1(e10, e11) = op1(e12, e11)) & ~ (op1(e10, e11) = op1(e11, e11)) & ~ (op1(e12, e10) = op1(e13, e10)) & ~ (op1(e11, e10) = op1(e13, e10)) & ~ (op1(e11, e10) = op1(e12, e10)) & ~ (op1(e10, e10) = op1(e13, e10)) & ~ (op1(e10, e10) = op1(e12, e10)) & ~ (op1(e10, e10) = op1(e11, e10))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG016+1.p', ax17)).
fof(f3544, plain, (spl12_25 | ~ spl12_6 | ~ spl12_18 | ~ spl12_37), inference(avatar_split_clause, [], [f3520, f690, f609, f558, f639])).
fof(f3520, plain, ((e10 = op1(e12, e11)) | (~ spl12_6 | ~ spl12_18 | ~ spl12_37)), inference(backward_demodulation, [], [f3499, f560])).
fof(f3499, plain, ((e10 = op1(e12, op1(e13, e12))) | (~ spl12_18 | ~ spl12_37)), inference(forward_demodulation, [], [f3491, f692])).
fof(f3491, plain, ((op1(e11, e12) = op1(e12, op1(e13, e12))) | ~ spl12_18), inference(backward_demodulation, [], [f124, f611])).
fof(f124, plain, (op1(op1(e12, e13), e12) = op1(e12, op1(e13, e12))), inference(cnf_transformation, [], [f2])).
fof(f3421, plain, (spl12_35 | ~ spl12_24 | ~ spl12_37 | ~ spl12_58 | ~ spl12_81), inference(avatar_split_clause, [], [f3412, f877, f779, f690, f634, f681])).
fof(f3412, plain, ((e12 = op1(e11, e13)) | (~ spl12_24 | ~ spl12_37 | ~ spl12_58 | ~ spl12_81)), inference(backward_demodulation, [], [f2996, f781])).
fof(f2996, plain, ((e12 = op1(op1(e10, e11), e13)) | (~ spl12_24 | ~ spl12_37 | ~ spl12_81)), inference(forward_demodulation, [], [f2985, f2504])).
fof(f2504, plain, ((e12 = op1(e10, op1(e10, e12))) | ~ spl12_81), inference(forward_demodulation, [], [f2489, f146])).
fof(f2489, plain, ((op1(e10, op1(e10, e12)) = op1(unit1, e12)) | ~ spl12_81), inference(backward_demodulation, [], [f80, f2478])).
fof(f80, plain, (op1(op1(e10, e10), e12) = op1(e10, op1(e10, e12))), inference(cnf_transformation, [], [f2])).
fof(f2985, plain, ((op1(e10, op1(e10, e12)) = op1(op1(e10, e11), e13)) | (~ spl12_24 | ~ spl12_37)), inference(backward_demodulation, [], [f85, f2984])).
fof(f85, plain, (op1(op1(e10, e11), e13) = op1(e10, op1(e11, e13))), inference(cnf_transformation, [], [f2])).
fof(f3420, plain, (spl12_11 | ~ spl12_24 | ~ spl12_25 | ~ spl12_37 | ~ spl12_58 | ~ spl12_81), inference(avatar_split_clause, [], [f3411, f877, f779, f690, f639, f634, f579])).
fof(f3411, plain, ((e12 = op1(e13, e11)) | (~ spl12_24 | ~ spl12_25 | ~ spl12_37 | ~ spl12_58 | ~ spl12_81)), inference(backward_demodulation, [], [f2979, f781])).
fof(f2979, plain, ((e12 = op1(e13, op1(e10, e11))) | (~ spl12_24 | ~ spl12_25 | ~ spl12_37 | ~ spl12_81)), inference(forward_demodulation, [], [f2968, f2879])).
fof(f2879, plain, ((e12 = op1(op1(e10, e12), e10)) | (~ spl12_25 | ~ spl12_37 | ~ spl12_81)), inference(forward_demodulation, [], [f2507, f2877])).
fof(f2877, plain, ((op1(e10, e12) = op1(e12, e10)) | (~ spl12_25 | ~ spl12_37)), inference(forward_demodulation, [], [f2732, f641])).
fof(f2507, plain, ((e12 = op1(op1(e12, e10), e10)) | ~ spl12_81), inference(forward_demodulation, [], [f2492, f147])).
fof(f2492, plain, ((op1(op1(e12, e10), e10) = op1(e12, unit1)) | ~ spl12_81), inference(backward_demodulation, [], [f110, f2478])).
fof(f110, plain, (op1(op1(e12, e10), e10) = op1(e12, op1(e10, e10))), inference(cnf_transformation, [], [f2])).
fof(f2968, plain, ((op1(op1(e10, e12), e10) = op1(e13, op1(e10, e11))) | (~ spl12_24 | ~ spl12_25 | ~ spl12_37)), inference(backward_demodulation, [], [f2905, f2958])).
fof(f2958, plain, ((op1(e10, e12) = op1(e13, e11)) | (~ spl12_24 | ~ spl12_25 | ~ spl12_37)), inference(forward_demodulation, [], [f2957, f2877])).
fof(f2957, plain, ((op1(e12, e10) = op1(e13, e11)) | (~ spl12_24 | ~ spl12_25)), inference(forward_demodulation, [], [f2956, f636])).
fof(f2956, plain, ((op1(e12, e10) = op1(op1(e12, e12), e11)) | ~ spl12_25), inference(forward_demodulation, [], [f119, f641])).
fof(f119, plain, (op1(op1(e12, e12), e11) = op1(e12, op1(e12, e11))), inference(cnf_transformation, [], [f2])).
fof(f2905, plain, ((op1(e13, op1(e10, e11)) = op1(op1(e13, e11), e10)) | (~ spl12_25 | ~ spl12_37)), inference(backward_demodulation, [], [f130, f2903])).
fof(f2903, plain, ((op1(e10, e11) = op1(e11, e10)) | (~ spl12_25 | ~ spl12_37)), inference(forward_demodulation, [], [f2730, f641])).
fof(f2730, plain, ((op1(e10, e11) = op1(e11, op1(e12, e11))) | ~ spl12_37), inference(backward_demodulation, [], [f103, f692])).
fof(f103, plain, (op1(op1(e11, e12), e11) = op1(e11, op1(e12, e11))), inference(cnf_transformation, [], [f2])).
fof(f130, plain, (op1(op1(e13, e11), e10) = op1(e13, op1(e11, e10))), inference(cnf_transformation, [], [f2])).
fof(f3383, plain, (~ spl12_37 | ~ spl12_66), inference(avatar_contradiction_clause, [], [f3382])).
fof(f3382, plain, ($false | (~ spl12_37 | ~ spl12_66)), inference(subsumption_resolution, [], [f3381, f415])).
fof(f415, plain, ~ (e10 = e12), inference(cnf_transformation, [], [f19])).
fof(f3381, plain, ((e10 = e12) | (~ spl12_37 | ~ spl12_66)), inference(forward_demodulation, [], [f3356, f692])).
fof(f3356, plain, ((e12 = op1(e11, e12)) | ~ spl12_66), inference(backward_demodulation, [], [f146, f815])).
fof(f815, plain, ((e11 = unit1) | ~ spl12_66), inference(avatar_component_clause, [], [f813])).
fof(f3334, plain, (~ spl12_90 | ~ spl12_102 | spl12_109 | ~ spl12_121), inference(avatar_contradiction_clause, [], [f3333])).
fof(f3333, plain, ($false | (~ spl12_90 | ~ spl12_102 | spl12_109 | ~ spl12_121)), inference(subsumption_resolution, [], [f3322, f997])).
fof(f997, plain, (~ (e20 = op2(e22, e21)) | spl12_109), inference(avatar_component_clause, [], [f996])).
fof(f3322, plain, ((e20 = op2(e22, e21)) | (~ spl12_90 | ~ spl12_102 | ~ spl12_121)), inference(backward_demodulation, [], [f3256, f917])).
fof(f3256, plain, ((e20 = op2(e22, op2(e23, e22))) | (~ spl12_102 | ~ spl12_121)), inference(forward_demodulation, [], [f3246, f1049])).
fof(f3246, plain, ((op2(e21, e22) = op2(e22, op2(e23, e22))) | ~ spl12_102), inference(backward_demodulation, [], [f225, f968])).
fof(f225, plain, (op2(op2(e22, e23), e22) = op2(e22, op2(e23, e22))), inference(cnf_transformation, [], [f6])).
fof(f3303, plain, (~ spl12_90 | ~ spl12_94), inference(avatar_split_clause, [], [f3290, f932, f915])).
fof(f932, plain, (spl12_94 <=> (e21 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl12_94])])).
fof(f3290, plain, (~ (e21 = op2(e23, e22)) | ~ spl12_94), inference(backward_demodulation, [], [f411, f934])).
fof(f934, plain, ((e21 = op2(e23, e21)) | ~ spl12_94), inference(avatar_component_clause, [], [f932])).
fof(f411, plain, ~ (op2(e23, e21) = op2(e23, e22)), inference(cnf_transformation, [], [f18])).
fof(f18, plain, (~ (op2(e23, e22) = op2(e23, e23)) & ~ (op2(e23, e21) = op2(e23, e23)) & ~ (op2(e23, e21) = op2(e23, e22)) & ~ (op2(e23, e20) = op2(e23, e23)) & ~ (op2(e23, e20) = op2(e23, e22)) & ~ (op2(e23, e20) = op2(e23, e21)) & ~ (op2(e22, e22) = op2(e22, e23)) & ~ (op2(e22, e21) = op2(e22, e23)) & ~ (op2(e22, e21) = op2(e22, e22)) & ~ (op2(e22, e20) = op2(e22, e23)) & ~ (op2(e22, e20) = op2(e22, e22)) & ~ (op2(e22, e20) = op2(e22, e21)) & ~ (op2(e21, e22) = op2(e21, e23)) & ~ (op2(e21, e21) = op2(e21, e23)) & ~ (op2(e21, e21) = op2(e21, e22)) & ~ (op2(e21, e20) = op2(e21, e23)) & ~ (op2(e21, e20) = op2(e21, e22)) & ~ (op2(e21, e20) = op2(e21, e21)) & ~ (op2(e20, e22) = op2(e20, e23)) & ~ (op2(e20, e21) = op2(e20, e23)) & ~ (op2(e20, e21) = op2(e20, e22)) & ~ (op2(e20, e20) = op2(e20, e23)) & ~ (op2(e20, e20) = op2(e20, e22)) & ~ (op2(e20, e20) = op2(e20, e21)) & ~ (op2(e22, e23) = op2(e23, e23)) & ~ (op2(e21, e23) = op2(e23, e23)) & ~ (op2(e21, e23) = op2(e22, e23)) & ~ (op2(e20, e23) = op2(e23, e23)) & ~ (op2(e20, e23) = op2(e22, e23)) & ~ (op2(e20, e23) = op2(e21, e23)) & ~ (op2(e22, e22) = op2(e23, e22)) & ~ (op2(e21, e22) = op2(e23, e22)) & ~ (op2(e21, e22) = op2(e22, e22)) & ~ (op2(e20, e22) = op2(e23, e22)) & ~ (op2(e20, e22) = op2(e22, e22)) & ~ (op2(e20, e22) = op2(e21, e22)) & ~ (op2(e22, e21) = op2(e23, e21)) & ~ (op2(e21, e21) = op2(e23, e21)) & ~ (op2(e21, e21) = op2(e22, e21)) & ~ (op2(e20, e21) = op2(e23, e21)) & ~ (op2(e20, e21) = op2(e22, e21)) & ~ (op2(e20, e21) = op2(e21, e21)) & ~ (op2(e22, e20) = op2(e23, e20)) & ~ (op2(e21, e20) = op2(e23, e20)) & ~ (op2(e21, e20) = op2(e22, e20)) & ~ (op2(e20, e20) = op2(e23, e20)) & ~ (op2(e20, e20) = op2(e22, e20)) & ~ (op2(e20, e20) = op2(e21, e20))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG016+1.p', ax18)).
fof(f3281, plain, (~ spl12_96 | ~ spl12_100), inference(avatar_split_clause, [], [f3267, f957, f940])).
fof(f940, plain, (spl12_96 <=> (e23 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl12_96])])).
fof(f3267, plain, (~ (e23 = op2(e23, e21)) | ~ spl12_100), inference(backward_demodulation, [], [f408, f959])).
fof(f408, plain, ~ (op2(e23, e20) = op2(e23, e21)), inference(cnf_transformation, [], [f18])).
fof(f3218, plain, (~ spl12_102 | ~ spl12_118), inference(avatar_split_clause, [], [f3208, f1034, f966])).
fof(f1034, plain, (spl12_118 <=> (e21 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl12_118])])).
fof(f3208, plain, (~ (e21 = op2(e22, e23)) | ~ spl12_118), inference(backward_demodulation, [], [f387, f1036])).
fof(f1036, plain, ((e21 = op2(e21, e23)) | ~ spl12_118), inference(avatar_component_clause, [], [f1034])).
fof(f387, plain, ~ (op2(e21, e23) = op2(e22, e23)), inference(cnf_transformation, [], [f18])).
fof(f3175, plain, (~ spl12_120 | ~ spl12_136), inference(avatar_split_clause, [], [f3165, f1110, f1042])).
fof(f1042, plain, (spl12_120 <=> (e23 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl12_120])])).
fof(f3165, plain, (~ (e23 = op2(e21, e23)) | ~ spl12_136), inference(backward_demodulation, [], [f384, f1112])).
fof(f384, plain, ~ (op2(e20, e23) = op2(e21, e23)), inference(cnf_transformation, [], [f18])).
fof(f3093, plain, (~ spl12_117 | ~ spl12_189), inference(avatar_split_clause, [], [f3081, f1478, f1030])).
fof(f1030, plain, (spl12_117 <=> (e20 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl12_117])])).
fof(f3081, plain, (~ (e20 = op2(e21, e23)) | ~ spl12_189), inference(backward_demodulation, [], [f1351, f1479])).
fof(f1351, plain, ~ (op2(e21, e23) = h3(e10)), inference(backward_demodulation, [], [f401, f1341])).
fof(f401, plain, ~ (op2(e21, e22) = op2(e21, e23)), inference(cnf_transformation, [], [f18])).
fof(f3061, plain, (spl12_165 | ~ spl12_149), inference(avatar_split_clause, [], [f2209, f1166, f1234])).
fof(f1234, plain, (spl12_165 <=> (e20 = inv2(e20))), introduced(avatar_definition, [new_symbols(naming, [spl12_165])])).
fof(f2209, plain, ((e20 = inv2(e20)) | ~ spl12_149), inference(backward_demodulation, [], [f291, f1168])).
fof(f3048, plain, (spl12_237 | ~ spl12_108), inference(avatar_split_clause, [], [f3047, f991, f1766])).
fof(f991, plain, (spl12_108 <=> (e23 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl12_108])])).
fof(f3047, plain, ((e23 = h3(e13)) | ~ spl12_108), inference(forward_demodulation, [], [f467, f993])).
fof(f993, plain, ((e23 = op2(e22, e22)) | ~ spl12_108), inference(avatar_component_clause, [], [f991])).
fof(f2951, plain, (~ spl12_56 | ~ spl12_24), inference(avatar_split_clause, [], [f2950, f634, f770])).
fof(f770, plain, (spl12_56 <=> (e13 = op1(e10, e12))), introduced(avatar_definition, [new_symbols(naming, [spl12_56])])).
fof(f2950, plain, (~ (e13 = op1(e10, e12)) | ~ spl12_24), inference(forward_demodulation, [], [f331, f636])).
fof(f331, plain, ~ (op1(e10, e12) = op1(e12, e12)), inference(cnf_transformation, [], [f17])).
fof(f2939, plain, (spl12_18 | ~ spl12_24), inference(avatar_split_clause, [], [f2938, f634, f609])).
fof(f2938, plain, ((e11 = op1(e12, e13)) | ~ spl12_24), inference(forward_demodulation, [], [f451, f636])).
fof(f451, plain, (e11 = op1(e12, op1(e12, e12))), inference(cnf_transformation, [], [f24])).
fof(f24, plain, ((e13 = op1(e12, e12)) & (e11 = op1(e12, op1(e12, e12))) & (e10 = op1(op1(e12, op1(e12, e12)), e12))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG016+1.p', ax24)).
fof(f2937, plain, (spl12_6 | ~ spl12_24), inference(avatar_split_clause, [], [f2936, f634, f558])).
fof(f2936, plain, ((e11 = op1(e13, e12)) | ~ spl12_24), inference(forward_demodulation, [], [f1283, f636])).
fof(f1283, plain, (e11 = op1(op1(e12, e12), e12)), inference(backward_demodulation, [], [f120, f451])).
fof(f120, plain, (op1(op1(e12, e12), e12) = op1(e12, op1(e12, e12))), inference(cnf_transformation, [], [f2])).
fof(f2894, plain, (~ spl12_53 | ~ spl12_37), inference(avatar_split_clause, [], [f2893, f690, f758])).
fof(f758, plain, (spl12_53 <=> (e10 = op1(e10, e12))), introduced(avatar_definition, [new_symbols(naming, [spl12_53])])).
fof(f2893, plain, (~ (e10 = op1(e10, e12)) | ~ spl12_37), inference(forward_demodulation, [], [f330, f692])).
fof(f330, plain, ~ (op1(e10, e12) = op1(e11, e12)), inference(cnf_transformation, [], [f17])).
fof(f2572, plain, ~ spl12_75, inference(avatar_contradiction_clause, [], [f2571])).
fof(f2571, plain, ($false | ~ spl12_75), inference(subsumption_resolution, [], [f2570, f417])).
fof(f2570, plain, ((e11 = e12) | ~ spl12_75), inference(backward_demodulation, [], [f146, f2560])).
fof(f2560, plain, ((e11 = op1(unit1, e12)) | ~ spl12_75), inference(backward_demodulation, [], [f1283, f2544])).
fof(f2544, plain, ((op1(e12, e12) = unit1) | ~ spl12_75), inference(backward_demodulation, [], [f155, f853])).
fof(f853, plain, ((e12 = inv1(e12)) | ~ spl12_75), inference(avatar_component_clause, [], [f851])).
fof(f2510, plain, (~ spl12_73 | ~ spl12_81), inference(avatar_split_clause, [], [f2482, f877, f843])).
fof(f843, plain, (spl12_73 <=> (e10 = inv1(e12))), introduced(avatar_definition, [new_symbols(naming, [spl12_73])])).
fof(f2482, plain, (~ (e10 = inv1(e12)) | ~ spl12_81), inference(backward_demodulation, [], [f286, f879])).
fof(f286, plain, ~ (inv1(e10) = inv1(e12)), inference(cnf_transformation, [], [f12])).
fof(f12, plain, (~ (inv1(e12) = inv1(e13)) & ~ (inv1(e11) = inv1(e13)) & ~ (inv1(e11) = inv1(e12)) & ~ (inv1(e10) = inv1(e13)) & ~ (inv1(e10) = inv1(e12)) & ~ (inv1(e10) = inv1(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG016+1.p', ax12)).
fof(f2397, plain, (~ spl12_93 | ~ spl12_109), inference(avatar_split_clause, [], [f2388, f996, f928])).
fof(f928, plain, (spl12_93 <=> (e20 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl12_93])])).
fof(f2388, plain, (~ (e20 = op2(e23, e21)) | ~ spl12_109), inference(backward_demodulation, [], [f377, f998])).
fof(f377, plain, ~ (op2(e22, e21) = op2(e23, e21)), inference(cnf_transformation, [], [f18])).
fof(f2345, plain, (spl12_189 | ~ spl12_121), inference(avatar_split_clause, [], [f2344, f1047, f1478])).
fof(f2344, plain, ((e20 = h3(e10)) | ~ spl12_121), inference(backward_demodulation, [], [f1341, f1049])).
fof(f2240, plain, (spl12_211 | ~ spl12_149 | ~ spl12_165), inference(avatar_split_clause, [], [f2211, f1234, f1166, f1588])).
fof(f2211, plain, ((e20 = h1(e13)) | (~ spl12_149 | ~ spl12_165)), inference(backward_demodulation, [], [f2080, f1168])).
fof(f2080, plain, ((unit2 = h1(e13)) | ~ spl12_165), inference(backward_demodulation, [], [f459, f2071])).
fof(f2071, plain, ((op2(e20, e20) = unit2) | ~ spl12_165), inference(backward_demodulation, [], [f252, f1236])).
fof(f1236, plain, ((e20 = inv2(e20)) | ~ spl12_165), inference(avatar_component_clause, [], [f1234])).
fof(f252, plain, (unit2 = op2(e20, inv2(e20))), inference(cnf_transformation, [], [f8])).
fof(f8, plain, (((e23 = inv2(e23)) | (e22 = inv2(e23)) | (e21 = inv2(e23)) | (e20 = inv2(e23))) & ((e23 = inv2(e22)) | (e22 = inv2(e22)) | (e21 = inv2(e22)) | (e20 = inv2(e22))) & ((e23 = inv2(e21)) | (e22 = inv2(e21)) | (e21 = inv2(e21)) | (e20 = inv2(e21))) & ((e23 = inv2(e20)) | (e22 = inv2(e20)) | (e21 = inv2(e20)) | (e20 = inv2(e20))) & (unit2 = op2(inv2(e23), e23)) & (unit2 = op2(e23, inv2(e23))) & (unit2 = op2(inv2(e22), e22)) & (unit2 = op2(e22, inv2(e22))) & (unit2 = op2(inv2(e21), e21)) & (unit2 = op2(e21, inv2(e21))) & (unit2 = op2(inv2(e20), e20)) & (unit2 = op2(e20, inv2(e20)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG016+1.p', ax8)).
fof(f2239, plain, (spl12_145 | ~ spl12_149 | ~ spl12_165), inference(avatar_split_clause, [], [f2210, f1234, f1166, f1149])).
fof(f2210, plain, ((e20 = op2(e20, e20)) | (~ spl12_149 | ~ spl12_165)), inference(backward_demodulation, [], [f2071, f1168])).
fof(f2237, plain, (spl12_136 | ~ spl12_149), inference(avatar_split_clause, [], [f2207, f1166, f1110])).
fof(f2207, plain, ((e23 = op2(e20, e23)) | ~ spl12_149), inference(backward_demodulation, [], [f249, f1168])).
fof(f249, plain, (e23 = op2(unit2, e23)), inference(cnf_transformation, [], [f7])).
fof(f2236, plain, (spl12_130 | ~ spl12_149), inference(avatar_split_clause, [], [f2206, f1166, f1085])).
fof(f2206, plain, ((e21 = op2(e21, e20)) | ~ spl12_149), inference(backward_demodulation, [], [f246, f1168])).
fof(f246, plain, (e21 = op2(e21, unit2)), inference(cnf_transformation, [], [f7])).
fof(f2173, plain, ~ spl12_159, inference(avatar_contradiction_clause, [], [f2172])).
fof(f2172, plain, ($false | ~ spl12_159), inference(subsumption_resolution, [], [f2171, f423])).
fof(f2171, plain, ((e21 = e22) | ~ spl12_159), inference(backward_demodulation, [], [f247, f2160])).
fof(f2160, plain, ((e21 = op2(unit2, e22)) | ~ spl12_159), inference(backward_demodulation, [], [f1337, f2148])).
fof(f2148, plain, ((unit2 = h3(e13)) | ~ spl12_159), inference(backward_demodulation, [], [f467, f2145])).
fof(f2145, plain, ((op2(e22, e22) = unit2) | ~ spl12_159), inference(backward_demodulation, [], [f256, f1210])).
fof(f1210, plain, ((e22 = inv2(e22)) | ~ spl12_159), inference(avatar_component_clause, [], [f1208])).
fof(f256, plain, (unit2 = op2(e22, inv2(e22))), inference(cnf_transformation, [], [f8])).
fof(f1817, plain, (~ spl12_233 | ~ spl12_234 | ~ spl12_235 | ~ spl12_236 | spl12_187 | spl12_184 | spl12_181 | ~ spl12_237 | ~ spl12_238 | ~ spl12_239 | ~ spl12_240 | ~ spl12_241 | ~ spl12_242 | ~ spl12_243 | ~ spl12_244 | ~ spl12_245 | ~ spl12_246 | ~ spl12_247 | ~ spl12_248 | ~ spl12_249), inference(avatar_split_clause, [], [f1748, f1814, f1810, f1806, f1802, f1798, f1794, f1790, f1786, f1782, f1778, f1774, f1770, f1766, f1433, f1450, f1467, f1762, f1758, f1754, f1750])).
fof(f1467, plain, (spl12_187 <=> sP6), introduced(avatar_definition, [new_symbols(naming, [spl12_187])])).
fof(f1450, plain, (spl12_184 <=> sP7), introduced(avatar_definition, [new_symbols(naming, [spl12_184])])).
fof(f1433, plain, (spl12_181 <=> sP8), introduced(avatar_definition, [new_symbols(naming, [spl12_181])])).
fof(f1748, plain, (~ (h3(op1(e10, e11)) = op2(h3(e10), e21)) | ~ (h3(op1(e10, e12)) = op2(e21, h3(e13))) | ~ (h3(op1(e11, e10)) = op2(e21, h3(e10))) | ~ (h2(e13) = h3(op1(e11, e11))) | ~ (h3(e10) = h3(op1(e11, e12))) | ~ (h3(op1(e11, e13)) = op2(e21, h3(e13))) | ~ (h3(op1(e12, e10)) = op2(e22, h3(e10))) | ~ (op2(e22, e21) = h3(op1(e12, e11))) | ~ (h3(e13) = h3(op1(e12, e12))) | ~ (e21 = h3(op1(e12, e13))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), e21)) | ~ (e21 = h3(op1(e13, e12))) | ~ (e23 = h3(e13)) | sP8 | sP7 | sP6 | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f1747, f1339])).
fof(f1747, plain, (~ (h3(op1(e10, e12)) = op2(e21, h3(e13))) | ~ (h3(op1(e11, e10)) = op2(e21, h3(e10))) | ~ (h2(e13) = h3(op1(e11, e11))) | ~ (h3(e10) = h3(op1(e11, e12))) | ~ (h3(op1(e11, e13)) = op2(e21, h3(e13))) | ~ (h3(op1(e12, e10)) = op2(e22, h3(e10))) | ~ (op2(e22, e21) = h3(op1(e12, e11))) | ~ (h3(e13) = h3(op1(e12, e12))) | ~ (e21 = h3(op1(e12, e13))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), e21)) | ~ (e21 = h3(op1(e13, e12))) | ~ (e23 = h3(e13)) | sP8 | sP7 | sP6 | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f1746, f1353])).
fof(f1353, plain, (op2(e21, h3(e13)) = op2(h3(e10), e22)), inference(backward_demodulation, [], [f1325, f1341])).
fof(f1325, plain, (op2(op2(e21, e22), e22) = op2(e21, h3(e13))), inference(backward_demodulation, [], [f205, f467])).
fof(f205, plain, (op2(op2(e21, e22), e22) = op2(e21, op2(e22, e22))), inference(cnf_transformation, [], [f6])).
fof(f1746, plain, (~ (h3(op1(e10, e12)) = op2(h3(e10), e22)) | ~ (h3(op1(e11, e10)) = op2(e21, h3(e10))) | ~ (h2(e13) = h3(op1(e11, e11))) | ~ (h3(e10) = h3(op1(e11, e12))) | ~ (h3(op1(e11, e13)) = op2(e21, h3(e13))) | ~ (h3(op1(e12, e10)) = op2(e22, h3(e10))) | ~ (op2(e22, e21) = h3(op1(e12, e11))) | ~ (h3(e13) = h3(op1(e12, e12))) | ~ (e21 = h3(op1(e12, e13))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), e21)) | ~ (e21 = h3(op1(e13, e12))) | ~ (e23 = h3(e13)) | sP8 | sP7 | sP6 | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f1745, f464])).
fof(f1745, plain, (~ (h3(op1(e11, e10)) = op2(e21, h3(e10))) | ~ (h2(e13) = h3(op1(e11, e11))) | ~ (h3(e10) = h3(op1(e11, e12))) | ~ (h3(op1(e11, e13)) = op2(e21, h3(e13))) | ~ (h3(op1(e12, e10)) = op2(e22, h3(e10))) | ~ (op2(e22, e21) = h3(op1(e12, e11))) | ~ (h3(e13) = h3(op1(e12, e12))) | ~ (e21 = h3(op1(e12, e13))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), e21)) | ~ (e21 = h3(op1(e13, e12))) | ~ (e23 = h3(e13)) | sP8 | sP7 | sP6 | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f1744, f1339])).
fof(f1744, plain, (~ (h2(e13) = h3(op1(e11, e11))) | ~ (h3(e10) = h3(op1(e11, e12))) | ~ (h3(op1(e11, e13)) = op2(e21, h3(e13))) | ~ (h3(op1(e12, e10)) = op2(e22, h3(e10))) | ~ (op2(e22, e21) = h3(op1(e12, e11))) | ~ (h3(e13) = h3(op1(e12, e12))) | ~ (e21 = h3(op1(e12, e13))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), e21)) | ~ (e21 = h3(op1(e13, e12))) | ~ (e23 = h3(e13)) | sP8 | sP7 | sP6 | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f1743, f463])).
fof(f1743, plain, (~ (op2(e21, e21) = h3(op1(e11, e11))) | ~ (h3(e10) = h3(op1(e11, e12))) | ~ (h3(op1(e11, e13)) = op2(e21, h3(e13))) | ~ (h3(op1(e12, e10)) = op2(e22, h3(e10))) | ~ (op2(e22, e21) = h3(op1(e12, e11))) | ~ (h3(e13) = h3(op1(e12, e12))) | ~ (e21 = h3(op1(e12, e13))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), e21)) | ~ (e21 = h3(op1(e13, e12))) | ~ (e23 = h3(e13)) | sP8 | sP7 | sP6 | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f1742, f1339])).
fof(f1742, plain, (~ (h3(e10) = h3(op1(e11, e12))) | ~ (h3(op1(e11, e13)) = op2(e21, h3(e13))) | ~ (h3(op1(e12, e10)) = op2(e22, h3(e10))) | ~ (op2(e22, e21) = h3(op1(e12, e11))) | ~ (h3(e13) = h3(op1(e12, e12))) | ~ (e21 = h3(op1(e12, e13))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), e21)) | ~ (e21 = h3(op1(e13, e12))) | ~ (e23 = h3(e13)) | sP8 | sP7 | sP6 | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f1741, f1341])).
fof(f1741, plain, (~ (op2(e21, e22) = h3(op1(e11, e12))) | ~ (h3(op1(e11, e13)) = op2(e21, h3(e13))) | ~ (h3(op1(e12, e10)) = op2(e22, h3(e10))) | ~ (op2(e22, e21) = h3(op1(e12, e11))) | ~ (h3(e13) = h3(op1(e12, e12))) | ~ (e21 = h3(op1(e12, e13))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), e21)) | ~ (e21 = h3(op1(e13, e12))) | ~ (e23 = h3(e13)) | sP8 | sP7 | sP6 | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f1740, f1339])).
fof(f1740, plain, (~ (h3(op1(e11, e12)) = op2(h3(e11), e22)) | ~ (h3(op1(e11, e13)) = op2(e21, h3(e13))) | ~ (h3(op1(e12, e10)) = op2(e22, h3(e10))) | ~ (op2(e22, e21) = h3(op1(e12, e11))) | ~ (h3(e13) = h3(op1(e12, e12))) | ~ (e21 = h3(op1(e12, e13))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), e21)) | ~ (e21 = h3(op1(e13, e12))) | ~ (e23 = h3(e13)) | sP8 | sP7 | sP6 | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f1739, f464])).
fof(f1739, plain, (~ (h3(op1(e11, e13)) = op2(e21, h3(e13))) | ~ (h3(op1(e12, e10)) = op2(e22, h3(e10))) | ~ (op2(e22, e21) = h3(op1(e12, e11))) | ~ (h3(e13) = h3(op1(e12, e12))) | ~ (e21 = h3(op1(e12, e13))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), e21)) | ~ (e21 = h3(op1(e13, e12))) | ~ (e23 = h3(e13)) | sP8 | sP7 | sP6 | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f1738, f1339])).
fof(f1738, plain, (~ (h3(op1(e12, e10)) = op2(e22, h3(e10))) | ~ (op2(e22, e21) = h3(op1(e12, e11))) | ~ (h3(e13) = h3(op1(e12, e12))) | ~ (e21 = h3(op1(e12, e13))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), e21)) | ~ (e21 = h3(op1(e13, e12))) | ~ (e23 = h3(e13)) | sP8 | sP7 | sP6 | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f1737, f464])).
fof(f1737, plain, (~ (op2(e22, e21) = h3(op1(e12, e11))) | ~ (h3(e13) = h3(op1(e12, e12))) | ~ (e21 = h3(op1(e12, e13))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), e21)) | ~ (e21 = h3(op1(e13, e12))) | ~ (e23 = h3(e13)) | sP8 | sP7 | sP6 | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f1736, f464])).
fof(f1736, plain, (~ (h3(op1(e12, e11)) = op2(h3(e12), e21)) | ~ (h3(e13) = h3(op1(e12, e12))) | ~ (e21 = h3(op1(e12, e13))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), e21)) | ~ (e21 = h3(op1(e13, e12))) | ~ (e23 = h3(e13)) | sP8 | sP7 | sP6 | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f1735, f1339])).
fof(f1735, plain, (~ (h3(e13) = h3(op1(e12, e12))) | ~ (e21 = h3(op1(e12, e13))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), e21)) | ~ (e21 = h3(op1(e13, e12))) | ~ (e23 = h3(e13)) | sP8 | sP7 | sP6 | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f1734, f467])).
fof(f1734, plain, (~ (op2(e22, e22) = h3(op1(e12, e12))) | ~ (e21 = h3(op1(e12, e13))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), e21)) | ~ (e21 = h3(op1(e13, e12))) | ~ (e23 = h3(e13)) | sP8 | sP7 | sP6 | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f1733, f464])).
fof(f1733, plain, (~ (e21 = h3(op1(e12, e13))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), e21)) | ~ (e21 = h3(op1(e13, e12))) | ~ (e23 = h3(e13)) | sP8 | sP7 | sP6 | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f1732, f1336])).
fof(f1732, plain, (~ (h3(op1(e12, e13)) = op2(e22, h3(e13))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), e21)) | ~ (e21 = h3(op1(e13, e12))) | ~ (e23 = h3(e13)) | sP8 | sP7 | sP6 | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f1731, f464])).
fof(f1731, plain, (~ (h3(op1(e13, e11)) = op2(h3(e13), e21)) | ~ (e21 = h3(op1(e13, e12))) | ~ (e23 = h3(e13)) | sP8 | sP7 | sP6 | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f1730, f1339])).
fof(f1730, plain, (~ (e21 = h3(op1(e13, e12))) | ~ (e23 = h3(e13)) | sP8 | sP7 | sP6 | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f1729, f1337])).
fof(f1729, plain, (~ (h3(op1(e13, e12)) = op2(h3(e13), e22)) | ~ (e23 = h3(e13)) | sP8 | sP7 | sP6 | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f531, f464])).
fof(f531, plain, (~ (e23 = h3(e13)) | sP8 | sP7 | sP6 | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(cnf_transformation, [], [f49])).
fof(f49, plain, (((~ (e23 = h4(e13)) & ~ (e23 = h4(e12)) & ~ (e23 = h4(e11)) & ~ (e23 = h4(e10))) | sP11 | sP10 | sP9 | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) & ((~ (e23 = h3(e13)) & ~ (e23 = h3(e12)) & ~ (e23 = h3(e11)) & ~ (e23 = h3(e10))) | sP8 | sP7 | sP6 | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) & ((~ (e23 = h2(e13)) & ~ (e23 = h2(e12)) & ~ (e23 = h2(e11)) & ~ (e23 = h2(e10))) | sP5 | sP4 | sP3 | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) & ((~ (e23 = h1(e13)) & ~ (e23 = h1(e12)) & ~ (e23 = h1(e11)) & ~ (e23 = h1(e10))) | sP2 | sP1 | sP0 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(definition_folding, [], [f36, e48, e47, e46, e45, e44, e43, e42, e41, e40, e39, e38, e37])).
fof(f37, plain, ((~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10))) | ~ sP0), inference(usedef, [], [e37])).
fof(e37, plain, (sP0 <=> (~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f38, plain, ((~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10))) | ~ sP1), inference(usedef, [], [e38])).
fof(e38, plain, (sP1 <=> (~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f39, plain, ((~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10))) | ~ sP2), inference(usedef, [], [e39])).
fof(e39, plain, (sP2 <=> (~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f40, plain, ((~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10))) | ~ sP3), inference(usedef, [], [e40])).
fof(e40, plain, (sP3 <=> (~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f41, plain, ((~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10))) | ~ sP4), inference(usedef, [], [e41])).
fof(e41, plain, (sP4 <=> (~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f42, plain, ((~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10))) | ~ sP5), inference(usedef, [], [e42])).
fof(e42, plain, (sP5 <=> (~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f43, plain, ((~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10))) | ~ sP6), inference(usedef, [], [e43])).
fof(e43, plain, (sP6 <=> (~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f44, plain, ((~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10))) | ~ sP7), inference(usedef, [], [e44])).
fof(e44, plain, (sP7 <=> (~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f45, plain, ((~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10))) | ~ sP8), inference(usedef, [], [e45])).
fof(e45, plain, (sP8 <=> (~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f46, plain, ((~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10))) | ~ sP9), inference(usedef, [], [e46])).
fof(e46, plain, (sP9 <=> (~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f47, plain, ((~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10))) | ~ sP10), inference(usedef, [], [e47])).
fof(e47, plain, (sP10 <=> (~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f48, plain, ((~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10))) | ~ sP11), inference(usedef, [], [e48])).
fof(e48, plain, (sP11 <=> (~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f36, plain, (((~ (e23 = h4(e13)) & ~ (e23 = h4(e12)) & ~ (e23 = h4(e11)) & ~ (e23 = h4(e10))) | (~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10))) | (~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10))) | (~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10))) | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) & ((~ (e23 = h3(e13)) & ~ (e23 = h3(e12)) & ~ (e23 = h3(e11)) & ~ (e23 = h3(e10))) | (~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10))) | (~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10))) | (~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10))) | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) & ((~ (e23 = h2(e13)) & ~ (e23 = h2(e12)) & ~ (e23 = h2(e11)) & ~ (e23 = h2(e10))) | (~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10))) | (~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10))) | (~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) & ((~ (e23 = h1(e13)) & ~ (e23 = h1(e12)) & ~ (e23 = h1(e11)) & ~ (e23 = h1(e10))) | (~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10))) | (~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10))) | (~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10))) | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(ennf_transformation, [], [f31])).
fof(f31, plain, ~ ((((e23 = h4(e13)) | (e23 = h4(e12)) | (e23 = h4(e11)) | (e23 = h4(e10))) & ((e22 = h4(e13)) | (e22 = h4(e12)) | (e22 = h4(e11)) | (e22 = h4(e10))) & ((e21 = h4(e13)) | (e21 = h4(e12)) | (e21 = h4(e11)) | (e21 = h4(e10))) & ((e20 = h4(e13)) | (e20 = h4(e12)) | (e20 = h4(e11)) | (e20 = h4(e10))) & (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) & (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) & (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) & (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) & (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) & (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) & (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) & (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) & (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) & (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) & (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) & (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) & (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) & (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) & (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) & (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) | (((e23 = h3(e13)) | (e23 = h3(e12)) | (e23 = h3(e11)) | (e23 = h3(e10))) & ((e22 = h3(e13)) | (e22 = h3(e12)) | (e22 = h3(e11)) | (e22 = h3(e10))) & ((e21 = h3(e13)) | (e21 = h3(e12)) | (e21 = h3(e11)) | (e21 = h3(e10))) & ((e20 = h3(e13)) | (e20 = h3(e12)) | (e20 = h3(e11)) | (e20 = h3(e10))) & (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) & (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) & (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) & (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) & (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) & (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) & (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) & (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) & (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) & (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) & (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) & (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) & (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) & (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) & (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) & (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) | (((e23 = h2(e13)) | (e23 = h2(e12)) | (e23 = h2(e11)) | (e23 = h2(e10))) & ((e22 = h2(e13)) | (e22 = h2(e12)) | (e22 = h2(e11)) | (e22 = h2(e10))) & ((e21 = h2(e13)) | (e21 = h2(e12)) | (e21 = h2(e11)) | (e21 = h2(e10))) & ((e20 = h2(e13)) | (e20 = h2(e12)) | (e20 = h2(e11)) | (e20 = h2(e10))) & (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) & (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) & (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) & (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) & (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) & (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) & (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) & (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) & (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) & (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) & (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) & (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) & (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) & (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) & (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) & (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) | (((e23 = h1(e13)) | (e23 = h1(e12)) | (e23 = h1(e11)) | (e23 = h1(e10))) & ((e22 = h1(e13)) | (e22 = h1(e12)) | (e22 = h1(e11)) | (e22 = h1(e10))) & ((e21 = h1(e13)) | (e21 = h1(e12)) | (e21 = h1(e11)) | (e21 = h1(e10))) & ((e20 = h1(e13)) | (e20 = h1(e12)) | (e20 = h1(e11)) | (e20 = h1(e10))) & (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) & (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) & (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) & (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) & (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) & (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) & (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) & (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) & (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) & (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) & (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) & (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) & (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) & (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) & (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) & (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(negated_conjecture, [], [f30])).
fof(f30, plain, ~ ((((e23 = h4(e13)) | (e23 = h4(e12)) | (e23 = h4(e11)) | (e23 = h4(e10))) & ((e22 = h4(e13)) | (e22 = h4(e12)) | (e22 = h4(e11)) | (e22 = h4(e10))) & ((e21 = h4(e13)) | (e21 = h4(e12)) | (e21 = h4(e11)) | (e21 = h4(e10))) & ((e20 = h4(e13)) | (e20 = h4(e12)) | (e20 = h4(e11)) | (e20 = h4(e10))) & (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) & (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) & (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) & (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) & (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) & (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) & (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) & (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) & (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) & (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) & (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) & (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) & (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) & (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) & (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) & (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) | (((e23 = h3(e13)) | (e23 = h3(e12)) | (e23 = h3(e11)) | (e23 = h3(e10))) & ((e22 = h3(e13)) | (e22 = h3(e12)) | (e22 = h3(e11)) | (e22 = h3(e10))) & ((e21 = h3(e13)) | (e21 = h3(e12)) | (e21 = h3(e11)) | (e21 = h3(e10))) & ((e20 = h3(e13)) | (e20 = h3(e12)) | (e20 = h3(e11)) | (e20 = h3(e10))) & (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) & (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) & (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) & (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) & (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) & (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) & (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) & (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) & (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) & (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) & (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) & (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) & (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) & (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) & (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) & (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) | (((e23 = h2(e13)) | (e23 = h2(e12)) | (e23 = h2(e11)) | (e23 = h2(e10))) & ((e22 = h2(e13)) | (e22 = h2(e12)) | (e22 = h2(e11)) | (e22 = h2(e10))) & ((e21 = h2(e13)) | (e21 = h2(e12)) | (e21 = h2(e11)) | (e21 = h2(e10))) & ((e20 = h2(e13)) | (e20 = h2(e12)) | (e20 = h2(e11)) | (e20 = h2(e10))) & (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) & (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) & (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) & (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) & (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) & (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) & (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) & (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) & (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) & (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) & (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) & (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) & (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) & (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) & (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) & (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) | (((e23 = h1(e13)) | (e23 = h1(e12)) | (e23 = h1(e11)) | (e23 = h1(e10))) & ((e22 = h1(e13)) | (e22 = h1(e12)) | (e22 = h1(e11)) | (e22 = h1(e10))) & ((e21 = h1(e13)) | (e21 = h1(e12)) | (e21 = h1(e11)) | (e21 = h1(e10))) & ((e20 = h1(e13)) | (e20 = h1(e12)) | (e20 = h1(e11)) | (e20 = h1(e10))) & (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) & (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) & (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) & (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) & (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) & (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) & (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) & (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) & (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) & (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) & (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) & (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) & (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) & (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) & (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) & (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG016+1.p', co1)).
fof(f1481, plain, (~ spl12_187 | ~ spl12_189), inference(avatar_split_clause, [], [f492, f1478, f1467])).
fof(f492, plain, (~ (e20 = h3(e10)) | ~ sP6), inference(cnf_transformation, [], [f55])).
fof(f55, plain, ((~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10))) | ~ sP6), inference(nnf_transformation, [], [f43])).
fof(f1460, plain, ~ spl12_184, inference(avatar_split_clause, [], [f1459, f1450])).
fof(f1459, plain, ~ sP7, inference(subsumption_resolution, [], [f489, f1339])).
fof(f489, plain, (~ (e21 = h3(e11)) | ~ sP7), inference(cnf_transformation, [], [f54])).
fof(f54, plain, ((~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10))) | ~ sP7), inference(nnf_transformation, [], [f44])).
fof(f1442, plain, ~ spl12_181, inference(avatar_split_clause, [], [f1441, f1433])).
fof(f1441, plain, ~ sP8, inference(subsumption_resolution, [], [f486, f464])).
fof(f486, plain, (~ (e22 = h3(e12)) | ~ sP8), inference(cnf_transformation, [], [f53])).
fof(f53, plain, ((~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10))) | ~ sP8), inference(nnf_transformation, [], [f45])).
fof(f1289, plain, spl12_121, inference(avatar_split_clause, [], [f1288, f1047])).
fof(f1288, plain, (e20 = op2(e21, e22)), inference(forward_demodulation, [], [f453, f454])).
fof(f453, plain, (e20 = op2(op2(e22, op2(e22, e22)), e22)), inference(cnf_transformation, [], [f25])).
fof(f1286, plain, spl12_108, inference(avatar_split_clause, [], [f455, f991])).
fof(f455, plain, (e23 = op2(e22, e22)), inference(cnf_transformation, [], [f25])).
fof(f1285, plain, spl12_37, inference(avatar_split_clause, [], [f1284, f690])).
fof(f1284, plain, (e10 = op1(e11, e12)), inference(forward_demodulation, [], [f450, f451])).
fof(f450, plain, (e10 = op1(op1(e12, op1(e12, e12)), e12)), inference(cnf_transformation, [], [f24])).
fof(f1282, plain, spl12_24, inference(avatar_split_clause, [], [f452, f634])).
fof(f452, plain, (e13 = op1(e12, e12)), inference(cnf_transformation, [], [f24])).
fof(f1254, plain, (~ spl12_74 | spl12_79), inference(avatar_split_clause, [], [f278, f868, f847])).
fof(f847, plain, (spl12_74 <=> (e11 = inv1(e12))), introduced(avatar_definition, [new_symbols(naming, [spl12_74])])).
fof(f278, plain, ((e12 = inv1(e11)) | ~ (e11 = inv1(e12))), inference(cnf_transformation, [], [f32])).
fof(f32, plain, (((e13 = inv1(e13)) | ~ (e13 = inv1(e13))) & ((e13 = inv1(e12)) | ~ (e12 = inv1(e13))) & ((e13 = inv1(e11)) | ~ (e11 = inv1(e13))) & ((e13 = inv1(e10)) | ~ (e10 = inv1(e13))) & ((e12 = inv1(e13)) | ~ (e13 = inv1(e12))) & ((e12 = inv1(e12)) | ~ (e12 = inv1(e12))) & ((e12 = inv1(e11)) | ~ (e11 = inv1(e12))) & ((e12 = inv1(e10)) | ~ (e10 = inv1(e12))) & ((e11 = inv1(e13)) | ~ (e13 = inv1(e11))) & ((e11 = inv1(e12)) | ~ (e12 = inv1(e11))) & ((e11 = inv1(e11)) | ~ (e11 = inv1(e11))) & ((e11 = inv1(e10)) | ~ (e10 = inv1(e11))) & ((e10 = inv1(e13)) | ~ (e13 = inv1(e10))) & ((e10 = inv1(e12)) | ~ (e12 = inv1(e10))) & ((e10 = inv1(e11)) | ~ (e11 = inv1(e10))) & ((e10 = inv1(e10)) | ~ (e10 = inv1(e10)))), inference(ennf_transformation, [], [f11])).
fof(f11, plain, (((e13 = inv1(e13)) => (e13 = inv1(e13))) & ((e12 = inv1(e13)) => (e13 = inv1(e12))) & ((e11 = inv1(e13)) => (e13 = inv1(e11))) & ((e10 = inv1(e13)) => (e13 = inv1(e10))) & ((e13 = inv1(e12)) => (e12 = inv1(e13))) & ((e12 = inv1(e12)) => (e12 = inv1(e12))) & ((e11 = inv1(e12)) => (e12 = inv1(e11))) & ((e10 = inv1(e12)) => (e12 = inv1(e10))) & ((e13 = inv1(e11)) => (e11 = inv1(e13))) & ((e12 = inv1(e11)) => (e11 = inv1(e12))) & ((e11 = inv1(e11)) => (e11 = inv1(e11))) & ((e10 = inv1(e11)) => (e11 = inv1(e10))) & ((e13 = inv1(e10)) => (e10 = inv1(e13))) & ((e12 = inv1(e10)) => (e10 = inv1(e12))) & ((e11 = inv1(e10)) => (e10 = inv1(e11))) & ((e10 = inv1(e10)) => (e10 = inv1(e10)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG016+1.p', ax11)).
fof(f1181, plain, (spl12_149 | spl12_150 | spl12_151 | spl12_152), inference(avatar_split_clause, [], [f251, f1178, f1174, f1170, f1166])).
fof(f251, plain, ((e23 = unit2) | (e22 = unit2) | (e21 = unit2) | (e20 = unit2)), inference(cnf_transformation, [], [f7])).
fof(f1045, plain, (spl12_117 | spl12_118 | spl12_119 | spl12_120), inference(avatar_split_clause, [], [f170, f1042, f1038, f1034, f1030])).
fof(f170, plain, ((e23 = op2(e21, e23)) | (e22 = op2(e21, e23)) | (e21 = op2(e21, e23)) | (e20 = op2(e21, e23))), inference(cnf_transformation, [], [f5])).
fof(f5, plain, (((e23 = op2(e23, e23)) | (e22 = op2(e23, e23)) | (e21 = op2(e23, e23)) | (e20 = op2(e23, e23))) & ((e23 = op2(e23, e22)) | (e22 = op2(e23, e22)) | (e21 = op2(e23, e22)) | (e20 = op2(e23, e22))) & ((e23 = op2(e23, e21)) | (e22 = op2(e23, e21)) | (e21 = op2(e23, e21)) | (e20 = op2(e23, e21))) & ((e23 = op2(e23, e20)) | (e22 = op2(e23, e20)) | (e21 = op2(e23, e20)) | (e20 = op2(e23, e20))) & ((e23 = op2(e22, e23)) | (e22 = op2(e22, e23)) | (e21 = op2(e22, e23)) | (e20 = op2(e22, e23))) & ((e23 = op2(e22, e22)) | (e22 = op2(e22, e22)) | (e21 = op2(e22, e22)) | (e20 = op2(e22, e22))) & ((e23 = op2(e22, e21)) | (e22 = op2(e22, e21)) | (e21 = op2(e22, e21)) | (e20 = op2(e22, e21))) & ((e23 = op2(e22, e20)) | (e22 = op2(e22, e20)) | (e21 = op2(e22, e20)) | (e20 = op2(e22, e20))) & ((e23 = op2(e21, e23)) | (e22 = op2(e21, e23)) | (e21 = op2(e21, e23)) | (e20 = op2(e21, e23))) & ((e23 = op2(e21, e22)) | (e22 = op2(e21, e22)) | (e21 = op2(e21, e22)) | (e20 = op2(e21, e22))) & ((e23 = op2(e21, e21)) | (e22 = op2(e21, e21)) | (e21 = op2(e21, e21)) | (e20 = op2(e21, e21))) & ((e23 = op2(e21, e20)) | (e22 = op2(e21, e20)) | (e21 = op2(e21, e20)) | (e20 = op2(e21, e20))) & ((e23 = op2(e20, e23)) | (e22 = op2(e20, e23)) | (e21 = op2(e20, e23)) | (e20 = op2(e20, e23))) & ((e23 = op2(e20, e22)) | (e22 = op2(e20, e22)) | (e21 = op2(e20, e22)) | (e20 = op2(e20, e22))) & ((e23 = op2(e20, e21)) | (e22 = op2(e20, e21)) | (e21 = op2(e20, e21)) | (e20 = op2(e20, e21))) & ((op2(e20, e20) = e23) | (op2(e20, e20) = e22) | (op2(e20, e20) = e21) | (e20 = op2(e20, e20)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG016+1.p', ax5)).
fof(f943, plain, (spl12_93 | spl12_94 | spl12_95 | spl12_96), inference(avatar_split_clause, [], [f176, f940, f936, f932, f928])).
fof(f176, plain, ((e23 = op2(e23, e21)) | (e22 = op2(e23, e21)) | (e21 = op2(e23, e21)) | (e20 = op2(e23, e21))), inference(cnf_transformation, [], [f5])).
fof(f858, plain, (spl12_73 | spl12_74 | spl12_75 | spl12_76), inference(avatar_split_clause, [], [f161, f855, f851, f847, f843])).
fof(f161, plain, ((e13 = inv1(e12)) | (e12 = inv1(e12)) | (e11 = inv1(e12)) | (e10 = inv1(e12))), inference(cnf_transformation, [], [f4])).
fof(f824, plain, (spl12_65 | spl12_66 | spl12_67 | spl12_68), inference(avatar_split_clause, [], [f150, f821, f817, f813, f809])).
fof(f150, plain, ((e13 = unit1) | (e12 = unit1) | (e11 = unit1) | (e10 = unit1)), inference(cnf_transformation, [], [f3])).
fof(f773, plain, (spl12_53 | spl12_54 | spl12_55 | spl12_56), inference(avatar_split_clause, [], [f64, f770, f766, f762, f758])).
fof(f64, plain, ((e13 = op1(e10, e12)) | (e12 = op1(e10, e12)) | (e11 = op1(e10, e12)) | (e10 = op1(e10, e12))), inference(cnf_transformation, [], [f1])).
fof(f1, plain, (((e13 = op1(e13, e13)) | (e12 = op1(e13, e13)) | (e11 = op1(e13, e13)) | (e10 = op1(e13, e13))) & ((e13 = op1(e13, e12)) | (e12 = op1(e13, e12)) | (e11 = op1(e13, e12)) | (e10 = op1(e13, e12))) & ((e13 = op1(e13, e11)) | (e12 = op1(e13, e11)) | (e11 = op1(e13, e11)) | (e10 = op1(e13, e11))) & ((e13 = op1(e13, e10)) | (e12 = op1(e13, e10)) | (e11 = op1(e13, e10)) | (e10 = op1(e13, e10))) & ((e13 = op1(e12, e13)) | (e12 = op1(e12, e13)) | (e11 = op1(e12, e13)) | (e10 = op1(e12, e13))) & ((e13 = op1(e12, e12)) | (e12 = op1(e12, e12)) | (e11 = op1(e12, e12)) | (e10 = op1(e12, e12))) & ((e13 = op1(e12, e11)) | (e12 = op1(e12, e11)) | (e11 = op1(e12, e11)) | (e10 = op1(e12, e11))) & ((e13 = op1(e12, e10)) | (e12 = op1(e12, e10)) | (e11 = op1(e12, e10)) | (e10 = op1(e12, e10))) & ((e13 = op1(e11, e13)) | (e12 = op1(e11, e13)) | (e11 = op1(e11, e13)) | (e10 = op1(e11, e13))) & ((e13 = op1(e11, e12)) | (e12 = op1(e11, e12)) | (e11 = op1(e11, e12)) | (e10 = op1(e11, e12))) & ((e13 = op1(e11, e11)) | (e12 = op1(e11, e11)) | (e11 = op1(e11, e11)) | (e10 = op1(e11, e11))) & ((e13 = op1(e11, e10)) | (e12 = op1(e11, e10)) | (e11 = op1(e11, e10)) | (e10 = op1(e11, e10))) & ((e13 = op1(e10, e13)) | (e12 = op1(e10, e13)) | (e11 = op1(e10, e13)) | (e10 = op1(e10, e13))) & ((e13 = op1(e10, e12)) | (e12 = op1(e10, e12)) | (e11 = op1(e10, e12)) | (e10 = op1(e10, e12))) & ((e13 = op1(e10, e11)) | (e12 = op1(e10, e11)) | (e11 = op1(e10, e11)) | (e10 = op1(e10, e11))) & ((op1(e10, e10) = e13) | (op1(e10, e10) = e12) | (op1(e10, e10) = e11) | (e10 = op1(e10, e10)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG016+1.p', ax1)).