fof(f6371, plain, $false, inference(avatar_sat_refutation, [], [f1257, f1278, f1299, f1320, f1362, f1383, f1404, f1425, f1467, f1488, f1509, f1530, f1572, f1593, f1614, f1635, f1764, f1765, f1766, f1767, f1768, f1769, f1771, f1774, f1775, f1776, f1777, f1778, f1779, f1780, f1781, f1782, f1783, f1786, f1787, f1788, f1789, f1790, f1791, f1792, f1793, f1794, f1795, f1798, f1799, f1802, f1803, f1804, f1805, f1806, f1807, f1810, f1811, f1812, f1813, f1828, f1834, f1842, f1849, f1855, f1856, f1862, f1863, f1870, f1877, f1883, f1884, f1891, f1897, f1898, f1905, f1912, f1918, f1919, f1925, f1926, f1933, f1940, f1946, f1954, f1961, f1968, f1975, f1981, f1982, f1989, f1996, f2003, f2010, f2017, f2023, f2030, f2031, f2038, f2045, f2051, f2052, f2059, f2065, f2066, f2073, f2079, f2080, f2087, f2094, f2101, f2107, f2115, f2122, f2129, f2135, f2136, f2143, f2149, f2150, f2157, f2164, f2170, f2171, f2178, f2185, f2192, f2199, f2206, f2212, f2219, f2220, f2226, f2227, f2234, f2241, f2248, f2255, f2262, f2268, f2276, f2283, f2290, f2297, f2303, f2304, f2311, f2317, f2318, f2325, f2332, f2338, f2339, f2345, f2346, f2353, f2360, f2367, f2374, f2381, f2388, f2395, f2401, f2409, f2416, f2423, f2429, f2437, f2560, f2569, f2578, f2600, f2609, f2618, f2640, f2649, f2658, f2680, f2689, f2707, f2713, f2718, f2725, f2731, f2737, f2743, f2748, f2755, f2761, f2767, f2773, f2778, f2785, f2791, f2797, f2803, f2808, f2815, f2821, f2827, f2833, f2838, f2844, f2846, f2847, f2853, f2854, f2856, f2862, f2863, f2864, f2866, f2875, f2888, f2901, f2914, f2993, f3028, f3037, f3100, f3104, f3105, f3108, f3109, f3110, f3115, f3130, f3144, f3145, f3157, f3160, f3172, f3174, f3185, f3194, f3203, f3220, f3222, f3223, f3224, f3234, f3235, f3237, f3238, f3239, f3241, f3250, f3251, f3252, f3253, f3254, f3255, f3256, f3265, f3266, f3267, f3268, f3269, f3270, f3278, f3279, f3280, f3281, f3282, f3290, f3291, f3292, f3293, f3304, f3305, f3306, f3308, f3309, f3310, f3331, f3362, f3370, f3371, f3373, f3374, f3375, f3386, f3403, f3419, f3426, f3428, f3429, f3430, f3456, f3457, f3460, f3473, f3474, f3475, f3497, f3500, f3513, f3518, f3530, f3583, f3595, f3596, f3603, f3622, f3633, f3656, f3662, f3666, f3684, f3697, f3700, f3708, f3726, f3732, f3765, f3788, f3789, f3816, f3837, f3863, f3872, f3878, f3881, f3887, f3912, f3917, f3935, f3942, f3965, f3969, f3972, f3978, f3981, f3999, f4013, f4026, f4029, f4036, f4040, f4044, f4045, f4046, f4047, f4052, f4055, f4060, f4066, f4071, f4087, f4093, f4124, f4157, f4172, f4188, f4225, f4229, f4237, f4262, f4284, f4302, f4329, f4345, f4371, f4382, f4386, f4387, f4390, f4406, f4416, f4455, f4458, f4463, f4468, f4475, f4480, f4491, f4517, f4518, f4523, f4538, f4553, f4557, f4558, f4582, f4596, f4605, f4631, f4638, f4639, f4670, f4699, f4706, f4730, f4731, f4766, f4769, f4771, f4802, f4806, f4818, f4837, f4844, f4863, f4864, f4899, f4948, f4949, f4956, f4957, f5001, f5008, f5032, f5048, f5059, f5069, f5085, f5086, f5091, f5102, f5106, f5150, f5177, f5184, f5218, f5222, f5228, f5245, f5257, f5259, f5275, f5278, f5283, f5289, f5291, f5305, f5308, f5312, f5313, f5319, f5321, f5332, f5341, f5370, f5397, f5411, f5416, f5417, f5418, f5428, f5430, f5434, f5443, f5452, f5468, f5476, f5482, f5498, f5510, f5511, f5546, f5550, f5558, f5571, f5585, f5592, f5596, f5597, f5609, f5610, f5616, f5628, f5639, f5640, f5644, f5661, f5690, f5704, f5714, f5720, f5749, f5751, f5791, f5797, f5804, f5833, f5835, f5843, f5854, f5867, f5879, f5880, f5897, f5913, f5929, f5936, f5955, f5961, f5972, f5982, f5984, f5986, f5988, f6000, f6001, f6009, f6011, f6018, f6022, f6029, f6038, f6044, f6047, f6054, f6062, f6076, f6079, f6082, f6084, f6092, f6103, f6105, f6107, f6116, f6124, f6126, f6128, f6130, f6137, f6138, f6150, f6151, f6165, f6167, f6179, f6180, f6189, f6202, f6203, f6242, f6271, f6279, f6288, f6297, f6304, f6305, f6312, f6321, f6325, f6328, f6337, f6348, f6360])).
fof(f6360, plain, (~ spl140_27 | ~ spl140_37), inference(avatar_split_clause, [], [f6351, f1389, f1347])).
fof(f1347, plain, (spl140_27 <=> (e1 = op(e3, e4))), introduced(avatar_definition, [new_symbols(naming, [spl140_27])])).
fof(f1389, plain, (spl140_37 <=> (e1 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl140_37])])).
fof(f6351, plain, (~ (e1 = op(e3, e4)) | ~ spl140_37), inference(backward_demodulation, [], [f466, f1391])).
fof(f1391, plain, ((e1 = op(e3, e2)) | ~ spl140_37), inference(avatar_component_clause, [], [f1389])).
fof(f466, plain, ~ (op(e3, e2) = op(e3, e4)), inference(cnf_transformation, [], [f5])).
fof(f5, plain, (~ (op(e4, e3) = op(e4, e4)) & ~ (op(e4, e2) = op(e4, e4)) & ~ (op(e4, e2) = op(e4, e3)) & ~ (op(e4, e1) = op(e4, e4)) & ~ (op(e4, e1) = op(e4, e3)) & ~ (op(e4, e1) = op(e4, e2)) & ~ (op(e4, e0) = op(e4, e4)) & ~ (op(e4, e0) = op(e4, e3)) & ~ (op(e4, e0) = op(e4, e2)) & ~ (op(e4, e0) = op(e4, e1)) & ~ (op(e3, e3) = op(e3, e4)) & ~ (op(e3, e2) = op(e3, e4)) & ~ (op(e3, e2) = op(e3, e3)) & ~ (op(e3, e1) = op(e3, e4)) & ~ (op(e3, e1) = op(e3, e3)) & ~ (op(e3, e1) = op(e3, e2)) & ~ (op(e3, e0) = op(e3, e4)) & ~ (op(e3, e0) = op(e3, e3)) & ~ (op(e3, e0) = op(e3, e2)) & ~ (op(e3, e0) = op(e3, e1)) & ~ (op(e2, e3) = op(e2, e4)) & ~ (op(e2, e2) = op(e2, e4)) & ~ (op(e2, e2) = op(e2, e3)) & ~ (op(e2, e1) = op(e2, e4)) & ~ (op(e2, e1) = op(e2, e3)) & ~ (op(e2, e1) = op(e2, e2)) & ~ (op(e2, e0) = op(e2, e4)) & ~ (op(e2, e0) = op(e2, e3)) & ~ (op(e2, e0) = op(e2, e2)) & ~ (op(e2, e0) = op(e2, e1)) & ~ (op(e1, e3) = op(e1, e4)) & ~ (op(e1, e2) = op(e1, e4)) & ~ (op(e1, e2) = op(e1, e3)) & ~ (op(e1, e1) = op(e1, e4)) & ~ (op(e1, e1) = op(e1, e3)) & ~ (op(e1, e1) = op(e1, e2)) & ~ (op(e1, e0) = op(e1, e4)) & ~ (op(e1, e0) = op(e1, e3)) & ~ (op(e1, e0) = op(e1, e2)) & ~ (op(e1, e0) = op(e1, e1)) & ~ (op(e0, e3) = op(e0, e4)) & ~ (op(e0, e2) = op(e0, e4)) & ~ (op(e0, e2) = op(e0, e3)) & ~ (op(e0, e1) = op(e0, e4)) & ~ (op(e0, e1) = op(e0, e3)) & ~ (op(e0, e1) = op(e0, e2)) & ~ (op(e0, e0) = op(e0, e4)) & ~ (op(e0, e0) = op(e0, e3)) & ~ (op(e0, e0) = op(e0, e2)) & ~ (op(e0, e0) = op(e0, e1)) & ~ (op(e3, e4) = op(e4, e4)) & ~ (op(e2, e4) = op(e4, e4)) & ~ (op(e2, e4) = op(e3, e4)) & ~ (op(e1, e4) = op(e4, e4)) & ~ (op(e1, e4) = op(e3, e4)) & ~ (op(e1, e4) = op(e2, e4)) & ~ (op(e0, e4) = op(e4, e4)) & ~ (op(e0, e4) = op(e3, e4)) & ~ (op(e0, e4) = op(e2, e4)) & ~ (op(e0, e4) = op(e1, e4)) & ~ (op(e3, e3) = op(e4, e3)) & ~ (op(e2, e3) = op(e4, e3)) & ~ (op(e2, e3) = op(e3, e3)) & ~ (op(e1, e3) = op(e4, e3)) & ~ (op(e1, e3) = op(e3, e3)) & ~ (op(e1, e3) = op(e2, e3)) & ~ (op(e0, e3) = op(e4, e3)) & ~ (op(e0, e3) = op(e3, e3)) & ~ (op(e0, e3) = op(e2, e3)) & ~ (op(e0, e3) = op(e1, e3)) & ~ (op(e3, e2) = op(e4, e2)) & ~ (op(e2, e2) = op(e4, e2)) & ~ (op(e2, e2) = op(e3, e2)) & ~ (op(e1, e2) = op(e4, e2)) & ~ (op(e1, e2) = op(e3, e2)) & ~ (op(e1, e2) = op(e2, e2)) & ~ (op(e0, e2) = op(e4, e2)) & ~ (op(e0, e2) = op(e3, e2)) & ~ (op(e0, e2) = op(e2, e2)) & ~ (op(e0, e2) = op(e1, e2)) & ~ (op(e3, e1) = op(e4, e1)) & ~ (op(e2, e1) = op(e4, e1)) & ~ (op(e2, e1) = op(e3, e1)) & ~ (op(e1, e1) = op(e4, e1)) & ~ (op(e1, e1) = op(e3, e1)) & ~ (op(e1, e1) = op(e2, e1)) & ~ (op(e0, e1) = op(e4, e1)) & ~ (op(e0, e1) = op(e3, e1)) & ~ (op(e0, e1) = op(e2, e1)) & ~ (op(e0, e1) = op(e1, e1)) & ~ (op(e3, e0) = op(e4, e0)) & ~ (op(e2, e0) = op(e4, e0)) & ~ (op(e2, e0) = op(e3, e0)) & ~ (op(e1, e0) = op(e4, e0)) & ~ (op(e1, e0) = op(e3, e0)) & ~ (op(e1, e0) = op(e2, e0)) & ~ (op(e0, e0) = op(e4, e0)) & ~ (op(e0, e0) = op(e3, e0)) & ~ (op(e0, e0) = op(e2, e0)) & ~ (op(e0, e0) = op(e1, e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG047+1.p', ax5)).
fof(f6348, plain, (~ spl140_40 | ~ spl140_45), inference(avatar_split_clause, [], [f6340, f1422, f1401])).
fof(f1401, plain, (spl140_40 <=> (e4 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl140_40])])).
fof(f1422, plain, (spl140_45 <=> (e4 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl140_45])])).
fof(f6340, plain, (~ (e4 = op(e3, e2)) | ~ spl140_45), inference(backward_demodulation, [], [f462, f1424])).
fof(f1424, plain, ((e4 = op(e3, e1)) | ~ spl140_45), inference(avatar_component_clause, [], [f1422])).
fof(f462, plain, ~ (op(e3, e1) = op(e3, e2)), inference(cnf_transformation, [], [f5])).
fof(f6337, plain, (~ spl140_27 | ~ spl140_52), inference(avatar_split_clause, [], [f6329, f1452, f1347])).
fof(f1452, plain, (spl140_52 <=> (e1 = op(e2, e4))), introduced(avatar_definition, [new_symbols(naming, [spl140_52])])).
fof(f6329, plain, (~ (e1 = op(e3, e4)) | ~ spl140_52), inference(backward_demodulation, [], [f425, f1454])).
fof(f1454, plain, ((e1 = op(e2, e4)) | ~ spl140_52), inference(avatar_component_clause, [], [f1452])).
fof(f425, plain, ~ (op(e2, e4) = op(e3, e4)), inference(cnf_transformation, [], [f5])).
fof(f6328, plain, (~ spl140_51 | ~ spl140_61), inference(avatar_split_clause, [], [f6327, f1490, f1448])).
fof(f1448, plain, (spl140_51 <=> (op(e2, e4) = unit)), introduced(avatar_definition, [new_symbols(naming, [spl140_51])])).
fof(f1490, plain, (spl140_61 <=> (op(e2, e2) = unit)), introduced(avatar_definition, [new_symbols(naming, [spl140_61])])).
fof(f6327, plain, (~ (op(e2, e4) = unit) | ~ spl140_61), inference(backward_demodulation, [], [f456, f1492])).
fof(f1492, plain, ((op(e2, e2) = unit) | ~ spl140_61), inference(avatar_component_clause, [], [f1490])).
fof(f456, plain, ~ (op(e2, e2) = op(e2, e4)), inference(cnf_transformation, [], [f5])).
fof(f6325, plain, (~ spl140_81 | ~ spl140_117 | ~ spl140_334), inference(avatar_contradiction_clause, [], [f6324])).
fof(f6324, plain, ($false | (~ spl140_81 | ~ spl140_117 | ~ spl140_334)), inference(subsumption_resolution, [], [f6323, f483])).
fof(f483, plain, ~ (e1 = e3), inference(cnf_transformation, [], [f6])).
fof(f6, plain, (~ (e3 = e4) & ~ (e2 = e4) & ~ (e2 = e3) & ~ (e1 = e4) & ~ (e1 = e3) & ~ (e1 = e2) & ~ (e0 = e4) & ~ (e0 = e3) & ~ (e0 = e2) & ~ (e0 = e1)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG047+1.p', ax6)).
fof(f6323, plain, ((e1 = e3) | (~ spl140_81 | ~ spl140_117 | ~ spl140_334)), inference(forward_demodulation, [], [f6320, f1727])).
fof(f1727, plain, ((e1 = op(unit, e1)) | ~ spl140_117), inference(avatar_component_clause, [], [f1725])).
fof(f1725, plain, (spl140_117 <=> (e1 = op(unit, e1))), introduced(avatar_definition, [new_symbols(naming, [spl140_117])])).
fof(f6320, plain, ((e3 = op(unit, e1)) | (~ spl140_81 | ~ spl140_334)), inference(backward_demodulation, [], [f3036, f1576])).
fof(f1576, plain, ((op(e1, e3) = unit) | ~ spl140_81), inference(avatar_component_clause, [], [f1574])).
fof(f1574, plain, (spl140_81 <=> (op(e1, e3) = unit)), introduced(avatar_definition, [new_symbols(naming, [spl140_81])])).
fof(f3036, plain, ((e3 = op(op(e1, e3), e1)) | ~ spl140_334), inference(avatar_component_clause, [], [f3034])).
fof(f3034, plain, (spl140_334 <=> (e3 = op(op(e1, e3), e1))), introduced(avatar_definition, [new_symbols(naming, [spl140_334])])).
fof(f6321, plain, (~ spl140_31 | ~ spl140_81), inference(avatar_split_clause, [], [f6315, f1574, f1364])).
fof(f1364, plain, (spl140_31 <=> (op(e3, e3) = unit)), introduced(avatar_definition, [new_symbols(naming, [spl140_31])])).
fof(f6315, plain, (~ (op(e3, e3) = unit) | ~ spl140_81), inference(backward_demodulation, [], [f413, f1576])).
fof(f413, plain, ~ (op(e1, e3) = op(e3, e3)), inference(cnf_transformation, [], [f5])).
fof(f6312, plain, (~ spl140_40 | ~ spl140_90), inference(avatar_split_clause, [], [f6307, f1611, f1401])).
fof(f1611, plain, (spl140_90 <=> (e4 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl140_90])])).
fof(f6307, plain, (~ (e4 = op(e3, e2)) | ~ spl140_90), inference(backward_demodulation, [], [f403, f1613])).
fof(f1613, plain, ((e4 = op(e1, e2)) | ~ spl140_90), inference(avatar_component_clause, [], [f1611])).
fof(f403, plain, ~ (op(e1, e2) = op(e3, e2)), inference(cnf_transformation, [], [f5])).
fof(f6305, plain, (~ spl140_83 | ~ spl140_93), inference(avatar_split_clause, [], [f6300, f1624, f1582])).
fof(f1582, plain, (spl140_83 <=> (e2 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl140_83])])).
fof(f1624, plain, (spl140_93 <=> (e2 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl140_93])])).
fof(f6300, plain, (~ (e2 = op(e1, e3)) | ~ spl140_93), inference(backward_demodulation, [], [f443, f1626])).
fof(f1626, plain, ((e2 = op(e1, e1)) | ~ spl140_93), inference(avatar_component_clause, [], [f1624])).
fof(f443, plain, ~ (op(e1, e1) = op(e1, e3)), inference(cnf_transformation, [], [f5])).
fof(f6304, plain, (~ spl140_43 | ~ spl140_93), inference(avatar_split_clause, [], [f6298, f1624, f1414])).
fof(f1414, plain, (spl140_43 <=> (e2 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl140_43])])).
fof(f6298, plain, (~ (e2 = op(e3, e1)) | ~ spl140_93), inference(backward_demodulation, [], [f393, f1626])).
fof(f393, plain, ~ (op(e1, e1) = op(e3, e1)), inference(cnf_transformation, [], [f5])).
fof(f6297, plain, (~ spl140_7 | spl140_19 | ~ spl140_237), inference(avatar_contradiction_clause, [], [f6296])).
fof(f6296, plain, ($false | (~ spl140_7 | spl140_19 | ~ spl140_237)), inference(subsumption_resolution, [], [f6295, f1314])).
fof(f1314, plain, (~ (e3 = op(e4, e1)) | spl140_19), inference(avatar_component_clause, [], [f1313])).
fof(f1313, plain, (spl140_19 <=> (e3 = op(e4, e1))), introduced(avatar_definition, [new_symbols(naming, [spl140_19])])).
fof(f6295, plain, ((e3 = op(e4, e1)) | (~ spl140_7 | ~ spl140_237)), inference(forward_demodulation, [], [f2555, f1265])).
fof(f1265, plain, ((e1 = op(e4, e3)) | ~ spl140_7), inference(avatar_component_clause, [], [f1263])).
fof(f1263, plain, (spl140_7 <=> (e1 = op(e4, e3))), introduced(avatar_definition, [new_symbols(naming, [spl140_7])])).
fof(f2555, plain, ((e3 = op(e4, op(e4, e3))) | ~ spl140_237), inference(avatar_component_clause, [], [f2553])).
fof(f2553, plain, (spl140_237 <=> (e3 = op(e4, op(e4, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl140_237])])).
fof(f6288, plain, (~ spl140_83 | ~ spl140_69 | spl140_260), inference(avatar_split_clause, [], [f6287, f2655, f1523, f1582])).
fof(f1523, plain, (spl140_69 <=> (e3 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl140_69])])).
fof(f2655, plain, (spl140_260 <=> (e2 = op(e1, op(e2, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl140_260])])).
fof(f6287, plain, (~ (e2 = op(e1, e3)) | (~ spl140_69 | spl140_260)), inference(forward_demodulation, [], [f2657, f1525])).
fof(f1525, plain, ((e3 = op(e2, e1)) | ~ spl140_69), inference(avatar_component_clause, [], [f1523])).
fof(f2657, plain, (~ (e2 = op(e1, op(e2, e1))) | spl140_260), inference(avatar_component_clause, [], [f2655])).
fof(f6279, plain, (~ spl140_14 | spl140_28 | ~ spl140_303), inference(avatar_contradiction_clause, [], [f6278])).
fof(f6278, plain, ($false | (~ spl140_14 | spl140_28 | ~ spl140_303)), inference(subsumption_resolution, [], [f6277, f1352])).
fof(f1352, plain, (~ (e2 = op(e3, e4)) | spl140_28), inference(avatar_component_clause, [], [f1351])).
fof(f1351, plain, (spl140_28 <=> (e2 = op(e3, e4))), introduced(avatar_definition, [new_symbols(naming, [spl140_28])])).
fof(f6277, plain, ((e2 = op(e3, e4)) | (~ spl140_14 | ~ spl140_303)), inference(forward_demodulation, [], [f2900, f1294])).
fof(f1294, plain, ((e3 = op(e4, e2)) | ~ spl140_14), inference(avatar_component_clause, [], [f1292])).
fof(f1292, plain, (spl140_14 <=> (e3 = op(e4, e2))), introduced(avatar_definition, [new_symbols(naming, [spl140_14])])).
fof(f2900, plain, ((e2 = op(op(e4, e2), e4)) | ~ spl140_303), inference(avatar_component_clause, [], [f2898])).
fof(f2898, plain, (spl140_303 <=> (e2 = op(op(e4, e2), e4))), introduced(avatar_definition, [new_symbols(naming, [spl140_303])])).
fof(f6271, plain, (~ spl140_27 | ~ spl140_79 | spl140_331), inference(avatar_split_clause, [], [f6270, f3021, f1565, f1347])).
fof(f1565, plain, (spl140_79 <=> (e3 = op(e1, e4))), introduced(avatar_definition, [new_symbols(naming, [spl140_79])])).
fof(f3021, plain, (spl140_331 <=> (e1 = op(op(e1, e4), e4))), introduced(avatar_definition, [new_symbols(naming, [spl140_331])])).
fof(f6270, plain, (~ (e1 = op(e3, e4)) | (~ spl140_79 | spl140_331)), inference(forward_demodulation, [], [f3023, f1567])).
fof(f1567, plain, ((e3 = op(e1, e4)) | ~ spl140_79), inference(avatar_component_clause, [], [f1565])).
fof(f3023, plain, (~ (e1 = op(op(e1, e4), e4)) | spl140_331), inference(avatar_component_clause, [], [f3021])).
fof(f6242, plain, (~ spl140_85 | ~ spl140_79 | spl140_352), inference(avatar_split_clause, [], [f6163, f3196, f1565, f1590])).
fof(f1590, plain, (spl140_85 <=> (e4 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl140_85])])).
fof(f3196, plain, (spl140_352 <=> (e4 = op(e1, op(e1, e4)))), introduced(avatar_definition, [new_symbols(naming, [spl140_352])])).
fof(f6163, plain, (~ (e4 = op(e1, e3)) | (~ spl140_79 | spl140_352)), inference(backward_demodulation, [], [f3197, f1567])).
fof(f3197, plain, (~ (e4 = op(e1, op(e1, e4))) | spl140_352), inference(avatar_component_clause, [], [f3196])).
fof(f6203, plain, (~ spl140_28 | ~ spl140_43), inference(avatar_split_clause, [], [f6194, f1414, f1351])).
fof(f6194, plain, (~ (e2 = op(e3, e4)) | ~ spl140_43), inference(backward_demodulation, [], [f464, f1416])).
fof(f1416, plain, ((e2 = op(e3, e1)) | ~ spl140_43), inference(avatar_component_clause, [], [f1414])).
fof(f464, plain, ~ (op(e3, e1) = op(e3, e4)), inference(cnf_transformation, [], [f5])).
fof(f6202, plain, (~ spl140_18 | ~ spl140_43), inference(avatar_split_clause, [], [f6192, f1414, f1309])).
fof(f1309, plain, (spl140_18 <=> (e2 = op(e4, e1))), introduced(avatar_definition, [new_symbols(naming, [spl140_18])])).
fof(f6192, plain, (~ (e2 = op(e4, e1)) | ~ spl140_43), inference(backward_demodulation, [], [f397, f1416])).
fof(f397, plain, ~ (op(e3, e1) = op(e4, e1)), inference(cnf_transformation, [], [f5])).
fof(f6189, plain, (spl140_11 | ~ spl140_51 | ~ spl140_301), inference(avatar_split_clause, [], [f6183, f2890, f1448, f1280])).
fof(f1280, plain, (spl140_11 <=> (op(e4, e2) = unit)), introduced(avatar_definition, [new_symbols(naming, [spl140_11])])).
fof(f2890, plain, (spl140_301 <=> (op(e2, e4) = op(e4, e2))), introduced(avatar_definition, [new_symbols(naming, [spl140_301])])).
fof(f6183, plain, ((op(e4, e2) = unit) | (~ spl140_51 | ~ spl140_301)), inference(backward_demodulation, [], [f2892, f1450])).
fof(f1450, plain, ((op(e2, e4) = unit) | ~ spl140_51), inference(avatar_component_clause, [], [f1448])).
fof(f2892, plain, ((op(e2, e4) = op(e4, e2)) | ~ spl140_301), inference(avatar_component_clause, [], [f2890])).
fof(f6180, plain, (~ spl140_52 | ~ spl140_62), inference(avatar_split_clause, [], [f6174, f1494, f1452])).
fof(f1494, plain, (spl140_62 <=> (e1 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl140_62])])).
fof(f6174, plain, (~ (e1 = op(e2, e4)) | ~ spl140_62), inference(backward_demodulation, [], [f456, f1496])).
fof(f1496, plain, ((e1 = op(e2, e2)) | ~ spl140_62), inference(avatar_component_clause, [], [f1494])).
fof(f6179, plain, (~ spl140_37 | ~ spl140_62), inference(avatar_split_clause, [], [f6172, f1494, f1389])).
fof(f6172, plain, (~ (e1 = op(e3, e2)) | ~ spl140_62), inference(backward_demodulation, [], [f405, f1496])).
fof(f405, plain, ~ (op(e2, e2) = op(e3, e2)), inference(cnf_transformation, [], [f5])).
fof(f6167, plain, (spl140_45 | ~ spl140_79 | ~ spl140_332), inference(avatar_split_clause, [], [f6162, f3025, f1565, f1422])).
fof(f3025, plain, (spl140_332 <=> (e4 = op(op(e1, e4), e1))), introduced(avatar_definition, [new_symbols(naming, [spl140_332])])).
fof(f6162, plain, ((e4 = op(e3, e1)) | (~ spl140_79 | ~ spl140_332)), inference(backward_demodulation, [], [f3027, f1567])).
fof(f3027, plain, ((e4 = op(op(e1, e4), e1)) | ~ spl140_332), inference(avatar_component_clause, [], [f3025])).
fof(f6165, plain, (~ spl140_4 | ~ spl140_79), inference(avatar_split_clause, [], [f6158, f1565, f1250])).
fof(f1250, plain, (spl140_4 <=> (e3 = op(e4, e4))), introduced(avatar_definition, [new_symbols(naming, [spl140_4])])).
fof(f6158, plain, (~ (e3 = op(e4, e4)) | ~ spl140_79), inference(backward_demodulation, [], [f424, f1567])).
fof(f424, plain, ~ (op(e1, e4) = op(e4, e4)), inference(cnf_transformation, [], [f5])).
fof(f6151, plain, (~ spl140_76 | ~ spl140_86), inference(avatar_split_clause, [], [f6142, f1595, f1553])).
fof(f1553, plain, (spl140_76 <=> (op(e1, e4) = unit)), introduced(avatar_definition, [new_symbols(naming, [spl140_76])])).
fof(f1595, plain, (spl140_86 <=> (op(e1, e2) = unit)), introduced(avatar_definition, [new_symbols(naming, [spl140_86])])).
fof(f6142, plain, (~ (op(e1, e4) = unit) | ~ spl140_86), inference(backward_demodulation, [], [f446, f1597])).
fof(f1597, plain, ((op(e1, e2) = unit) | ~ spl140_86), inference(avatar_component_clause, [], [f1595])).
fof(f446, plain, ~ (op(e1, e2) = op(e1, e4)), inference(cnf_transformation, [], [f5])).
fof(f6150, plain, (~ spl140_11 | ~ spl140_86), inference(avatar_split_clause, [], [f6141, f1595, f1280])).
fof(f6141, plain, (~ (op(e4, e2) = unit) | ~ spl140_86), inference(backward_demodulation, [], [f404, f1597])).
fof(f404, plain, ~ (op(e1, e2) = op(e4, e2)), inference(cnf_transformation, [], [f5])).
fof(f6138, plain, (~ spl140_90 | ~ spl140_95), inference(avatar_split_clause, [], [f6134, f1632, f1611])).
fof(f1632, plain, (spl140_95 <=> (e4 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl140_95])])).
fof(f6134, plain, (~ (e4 = op(e1, e2)) | ~ spl140_95), inference(backward_demodulation, [], [f442, f1634])).
fof(f1634, plain, ((e4 = op(e1, e1)) | ~ spl140_95), inference(avatar_component_clause, [], [f1632])).
fof(f442, plain, ~ (op(e1, e1) = op(e1, e2)), inference(cnf_transformation, [], [f5])).
fof(f6137, plain, (~ spl140_45 | ~ spl140_95), inference(avatar_split_clause, [], [f6132, f1632, f1422])).
fof(f6132, plain, (~ (e4 = op(e3, e1)) | ~ spl140_95), inference(backward_demodulation, [], [f393, f1634])).
fof(f6130, plain, (~ spl140_45 | ~ spl140_7 | spl140_238), inference(avatar_split_clause, [], [f6129, f2557, f1263, f1422])).
fof(f2557, plain, (spl140_238 <=> (e4 = op(e3, op(e4, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl140_238])])).
fof(f6129, plain, (~ (e4 = op(e3, e1)) | (~ spl140_7 | spl140_238)), inference(forward_demodulation, [], [f2559, f1265])).
fof(f2559, plain, (~ (e4 = op(e3, op(e4, e3))) | spl140_238), inference(avatar_component_clause, [], [f2557])).
fof(f6128, plain, (~ spl140_28 | ~ spl140_60 | spl140_240), inference(avatar_split_clause, [], [f6127, f2566, f1485, f1351])).
fof(f1485, plain, (spl140_60 <=> (e4 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl140_60])])).
fof(f2566, plain, (spl140_240 <=> (e2 = op(e3, op(e2, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl140_240])])).
fof(f6127, plain, (~ (e2 = op(e3, e4)) | (~ spl140_60 | spl140_240)), inference(forward_demodulation, [], [f2568, f1487])).
fof(f1487, plain, ((e4 = op(e2, e3)) | ~ spl140_60), inference(avatar_component_clause, [], [f1485])).
fof(f2568, plain, (~ (e2 = op(e3, op(e2, e3))) | spl140_240), inference(avatar_component_clause, [], [f2566])).
fof(f6126, plain, (~ spl140_37 | ~ spl140_83 | spl140_242), inference(avatar_split_clause, [], [f6125, f2575, f1582, f1389])).
fof(f2575, plain, (spl140_242 <=> (e1 = op(e3, op(e1, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl140_242])])).
fof(f6125, plain, (~ (e1 = op(e3, e2)) | (~ spl140_83 | spl140_242)), inference(forward_demodulation, [], [f2577, f1584])).
fof(f1584, plain, ((e2 = op(e1, e3)) | ~ spl140_83), inference(avatar_component_clause, [], [f1582])).
fof(f2577, plain, (~ (e1 = op(e3, op(e1, e3))) | spl140_242), inference(avatar_component_clause, [], [f2575])).
fof(f6124, plain, (~ spl140_247 | ~ spl140_301 | spl140_350), inference(avatar_contradiction_clause, [], [f6123])).
fof(f6123, plain, ($false | (~ spl140_247 | ~ spl140_301 | spl140_350)), inference(subsumption_resolution, [], [f6122, f3188])).
fof(f3188, plain, (~ (e4 = op(e2, op(e2, e4))) | spl140_350), inference(avatar_component_clause, [], [f3187])).
fof(f3187, plain, (spl140_350 <=> (e4 = op(e2, op(e2, e4)))), introduced(avatar_definition, [new_symbols(naming, [spl140_350])])).
fof(f6122, plain, ((e4 = op(e2, op(e2, e4))) | (~ spl140_247 | ~ spl140_301)), inference(forward_demodulation, [], [f2598, f2892])).
fof(f2598, plain, ((e4 = op(e2, op(e4, e2))) | ~ spl140_247), inference(avatar_component_clause, [], [f2597])).
fof(f2597, plain, (spl140_247 <=> (e4 = op(e2, op(e4, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl140_247])])).
fof(f6116, plain, (spl140_79 | ~ spl140_7 | ~ spl140_300), inference(avatar_split_clause, [], [f6115, f2885, f1263, f1565])).
fof(f2885, plain, (spl140_300 <=> (e3 = op(op(e4, e3), e4))), introduced(avatar_definition, [new_symbols(naming, [spl140_300])])).
fof(f6115, plain, ((e3 = op(e1, e4)) | (~ spl140_7 | ~ spl140_300)), inference(forward_demodulation, [], [f2887, f1265])).
fof(f2887, plain, ((e3 = op(op(e4, e3), e4)) | ~ spl140_300), inference(avatar_component_clause, [], [f2885])).
fof(f6107, plain, (spl140_40 | ~ spl140_60 | ~ spl140_312), inference(avatar_split_clause, [], [f6106, f2938, f1485, f1401])).
fof(f2938, plain, (spl140_312 <=> (op(e2, e3) = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl140_312])])).
fof(f6106, plain, ((e4 = op(e3, e2)) | (~ spl140_60 | ~ spl140_312)), inference(forward_demodulation, [], [f2940, f1487])).
fof(f2940, plain, ((op(e2, e3) = op(e3, e2)) | ~ spl140_312), inference(avatar_component_clause, [], [f2938])).
fof(f6105, plain, (spl140_43 | ~ spl140_83 | ~ spl140_315), inference(avatar_split_clause, [], [f6104, f2951, f1582, f1414])).
fof(f2951, plain, (spl140_315 <=> (op(e1, e3) = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl140_315])])).
fof(f6104, plain, ((e2 = op(e3, e1)) | (~ spl140_83 | ~ spl140_315)), inference(forward_demodulation, [], [f2953, f1584])).
fof(f2953, plain, ((op(e1, e3) = op(e3, e1)) | ~ spl140_315), inference(avatar_component_clause, [], [f2951])).
fof(f6103, plain, (spl140_14 | ~ spl140_60 | ~ spl140_324), inference(avatar_split_clause, [], [f6102, f2990, f1485, f1292])).
fof(f2990, plain, (spl140_324 <=> (e3 = op(op(e2, e3), e2))), introduced(avatar_definition, [new_symbols(naming, [spl140_324])])).
fof(f6102, plain, ((e3 = op(e4, e2)) | (~ spl140_60 | ~ spl140_324)), inference(forward_demodulation, [], [f2992, f1487])).
fof(f2992, plain, ((e3 = op(op(e2, e3), e2)) | ~ spl140_324), inference(avatar_component_clause, [], [f2990])).
fof(f6092, plain, (spl140_69 | ~ spl140_83 | ~ spl140_334), inference(avatar_split_clause, [], [f6091, f3034, f1582, f1523])).
fof(f6091, plain, ((e3 = op(e2, e1)) | (~ spl140_83 | ~ spl140_334)), inference(forward_demodulation, [], [f3036, f1584])).
fof(f6084, plain, (~ spl140_36 | ~ spl140_31), inference(avatar_split_clause, [], [f6083, f1364, f1385])).
fof(f1385, plain, (spl140_36 <=> (op(e3, e2) = unit)), introduced(avatar_definition, [new_symbols(naming, [spl140_36])])).
fof(f6083, plain, (~ (op(e3, e2) = unit) | ~ spl140_31), inference(forward_demodulation, [], [f465, f1366])).
fof(f1366, plain, ((op(e3, e3) = unit) | ~ spl140_31), inference(avatar_component_clause, [], [f1364])).
fof(f465, plain, ~ (op(e3, e2) = op(e3, e3)), inference(cnf_transformation, [], [f5])).
fof(f6082, plain, (~ spl140_26 | ~ spl140_31), inference(avatar_split_clause, [], [f6081, f1364, f1343])).
fof(f1343, plain, (spl140_26 <=> (op(e3, e4) = unit)), introduced(avatar_definition, [new_symbols(naming, [spl140_26])])).
fof(f6081, plain, (~ (op(e3, e4) = unit) | ~ spl140_31), inference(forward_demodulation, [], [f467, f1366])).
fof(f467, plain, ~ (op(e3, e3) = op(e3, e4)), inference(cnf_transformation, [], [f5])).
fof(f6079, plain, (~ spl140_52 | spl140_70 | ~ spl140_350), inference(avatar_contradiction_clause, [], [f6078])).
fof(f6078, plain, ($false | (~ spl140_52 | spl140_70 | ~ spl140_350)), inference(subsumption_resolution, [], [f6077, f1528])).
fof(f1528, plain, (~ (e4 = op(e2, e1)) | spl140_70), inference(avatar_component_clause, [], [f1527])).
fof(f1527, plain, (spl140_70 <=> (e4 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl140_70])])).
fof(f6077, plain, ((e4 = op(e2, e1)) | (~ spl140_52 | ~ spl140_350)), inference(forward_demodulation, [], [f3189, f1454])).
fof(f3189, plain, ((e4 = op(e2, op(e2, e4))) | ~ spl140_350), inference(avatar_component_clause, [], [f3187])).
fof(f6076, plain, (spl140_78 | ~ spl140_90 | ~ spl140_250), inference(avatar_contradiction_clause, [], [f6075])).
fof(f6075, plain, ($false | (spl140_78 | ~ spl140_90 | ~ spl140_250)), inference(subsumption_resolution, [], [f6074, f1562])).
fof(f1562, plain, (~ (e2 = op(e1, e4)) | spl140_78), inference(avatar_component_clause, [], [f1561])).
fof(f1561, plain, (spl140_78 <=> (e2 = op(e1, e4))), introduced(avatar_definition, [new_symbols(naming, [spl140_78])])).
fof(f6074, plain, ((e2 = op(e1, e4)) | (~ spl140_90 | ~ spl140_250)), inference(forward_demodulation, [], [f2613, f1613])).
fof(f2613, plain, ((e2 = op(e1, op(e1, e2))) | ~ spl140_250), inference(avatar_component_clause, [], [f2611])).
fof(f2611, plain, (spl140_250 <=> (e2 = op(e1, op(e1, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl140_250])])).
fof(f6062, plain, (~ spl140_1 | ~ spl140_11), inference(avatar_split_clause, [], [f6057, f1280, f1238])).
fof(f1238, plain, (spl140_1 <=> (op(e4, e4) = unit)), introduced(avatar_definition, [new_symbols(naming, [spl140_1])])).
fof(f6057, plain, (~ (op(e4, e4) = unit) | ~ spl140_11), inference(backward_demodulation, [], [f476, f1282])).
fof(f1282, plain, ((op(e4, e2) = unit) | ~ spl140_11), inference(avatar_component_clause, [], [f1280])).
fof(f476, plain, ~ (op(e4, e2) = op(e4, e4)), inference(cnf_transformation, [], [f5])).
fof(f6054, plain, (~ spl140_35 | ~ spl140_45), inference(avatar_split_clause, [], [f6049, f1422, f1380])).
fof(f1380, plain, (spl140_35 <=> (e4 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl140_35])])).
fof(f6049, plain, (~ (e4 = op(e3, e3)) | ~ spl140_45), inference(backward_demodulation, [], [f463, f1424])).
fof(f463, plain, ~ (op(e3, e1) = op(e3, e3)), inference(cnf_transformation, [], [f5])).
fof(f6047, plain, (~ spl140_14 | ~ spl140_60 | spl140_324), inference(avatar_split_clause, [], [f6043, f2990, f1485, f1292])).
fof(f6043, plain, (~ (e3 = op(e4, e2)) | (~ spl140_60 | spl140_324)), inference(backward_demodulation, [], [f2991, f1487])).
fof(f2991, plain, (~ (e3 = op(op(e2, e3), e2)) | spl140_324), inference(avatar_component_clause, [], [f2990])).
fof(f6044, plain, (~ spl140_35 | ~ spl140_60), inference(avatar_split_clause, [], [f6039, f1485, f1380])).
fof(f6039, plain, (~ (e4 = op(e3, e3)) | ~ spl140_60), inference(backward_demodulation, [], [f415, f1487])).
fof(f415, plain, ~ (op(e2, e3) = op(e3, e3)), inference(cnf_transformation, [], [f5])).
fof(f6038, plain, (~ spl140_14 | ~ spl140_64), inference(avatar_split_clause, [], [f6034, f1502, f1292])).
fof(f1502, plain, (spl140_64 <=> (e3 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl140_64])])).
fof(f6034, plain, (~ (e3 = op(e4, e2)) | ~ spl140_64), inference(backward_demodulation, [], [f406, f1504])).
fof(f1504, plain, ((e3 = op(e2, e2)) | ~ spl140_64), inference(avatar_component_clause, [], [f1502])).
fof(f406, plain, ~ (op(e2, e2) = op(e4, e2)), inference(cnf_transformation, [], [f5])).
fof(f6029, plain, (~ spl140_41 | ~ spl140_66), inference(avatar_split_clause, [], [f6023, f1511, f1406])).
fof(f1406, plain, (spl140_41 <=> (op(e3, e1) = unit)), introduced(avatar_definition, [new_symbols(naming, [spl140_41])])).
fof(f1511, plain, (spl140_66 <=> (op(e2, e1) = unit)), introduced(avatar_definition, [new_symbols(naming, [spl140_66])])).
fof(f6023, plain, (~ (op(e3, e1) = unit) | ~ spl140_66), inference(backward_demodulation, [], [f395, f1513])).
fof(f1513, plain, ((op(e2, e1) = unit) | ~ spl140_66), inference(avatar_component_clause, [], [f1511])).
fof(f395, plain, ~ (op(e2, e1) = op(e3, e1)), inference(cnf_transformation, [], [f5])).
fof(f6022, plain, (spl140_22 | ~ spl140_76 | ~ spl140_353), inference(avatar_contradiction_clause, [], [f6021])).
fof(f6021, plain, ($false | (spl140_22 | ~ spl140_76 | ~ spl140_353)), inference(subsumption_resolution, [], [f6017, f1327])).
fof(f1327, plain, (~ (e1 = op(e4, unit)) | spl140_22), inference(avatar_component_clause, [], [f1326])).
fof(f1326, plain, (spl140_22 <=> (e1 = op(e4, unit))), introduced(avatar_definition, [new_symbols(naming, [spl140_22])])).
fof(f6017, plain, ((e1 = op(e4, unit)) | (~ spl140_76 | ~ spl140_353)), inference(backward_demodulation, [], [f3201, f1555])).
fof(f1555, plain, ((op(e1, e4) = unit) | ~ spl140_76), inference(avatar_component_clause, [], [f1553])).
fof(f3201, plain, ((e1 = op(e4, op(e1, e4))) | ~ spl140_353), inference(avatar_component_clause, [], [f3200])).
fof(f3200, plain, (spl140_353 <=> (e1 = op(e4, op(e1, e4)))), introduced(avatar_definition, [new_symbols(naming, [spl140_353])])).
fof(f6018, plain, (~ spl140_1 | ~ spl140_76), inference(avatar_split_clause, [], [f6013, f1553, f1238])).
fof(f6013, plain, (~ (op(e4, e4) = unit) | ~ spl140_76), inference(backward_demodulation, [], [f424, f1555])).
fof(f6011, plain, (~ spl140_52 | ~ spl140_90 | spl140_251), inference(avatar_contradiction_clause, [], [f6010])).
fof(f6010, plain, ($false | (~ spl140_52 | ~ spl140_90 | spl140_251)), inference(subsumption_resolution, [], [f6006, f1454])).
fof(f6006, plain, (~ (e1 = op(e2, e4)) | (~ spl140_90 | spl140_251)), inference(backward_demodulation, [], [f2617, f1613])).
fof(f2617, plain, (~ (e1 = op(e2, op(e1, e2))) | spl140_251), inference(avatar_component_clause, [], [f2615])).
fof(f2615, plain, (spl140_251 <=> (e1 = op(e2, op(e1, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl140_251])])).
fof(f6009, plain, (~ spl140_65 | ~ spl140_90), inference(avatar_split_clause, [], [f6002, f1611, f1506])).
fof(f1506, plain, (spl140_65 <=> (e4 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl140_65])])).
fof(f6002, plain, (~ (e4 = op(e2, e2)) | ~ spl140_90), inference(backward_demodulation, [], [f402, f1613])).
fof(f402, plain, ~ (op(e1, e2) = op(e2, e2)), inference(cnf_transformation, [], [f5])).
fof(f6001, plain, (~ spl140_79 | ~ spl140_94), inference(avatar_split_clause, [], [f5997, f1628, f1565])).
fof(f1628, plain, (spl140_94 <=> (e3 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl140_94])])).
fof(f5997, plain, (~ (e3 = op(e1, e4)) | ~ spl140_94), inference(backward_demodulation, [], [f444, f1630])).
fof(f1630, plain, ((e3 = op(e1, e1)) | ~ spl140_94), inference(avatar_component_clause, [], [f1628])).
fof(f444, plain, ~ (op(e1, e1) = op(e1, e4)), inference(cnf_transformation, [], [f5])).
fof(f6000, plain, (~ spl140_69 | ~ spl140_94), inference(avatar_split_clause, [], [f5994, f1628, f1523])).
fof(f5994, plain, (~ (e3 = op(e2, e1)) | ~ spl140_94), inference(backward_demodulation, [], [f392, f1630])).
fof(f392, plain, ~ (op(e1, e1) = op(e2, e1)), inference(cnf_transformation, [], [f5])).
fof(f5988, plain, (~ spl140_69 | ~ spl140_37 | spl140_249), inference(avatar_split_clause, [], [f5987, f2606, f1389, f1523])).
fof(f2606, plain, (spl140_249 <=> (e3 = op(e2, op(e3, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl140_249])])).
fof(f5987, plain, (~ (e3 = op(e2, e1)) | (~ spl140_37 | spl140_249)), inference(forward_demodulation, [], [f2608, f1391])).
fof(f2608, plain, (~ (e3 = op(e2, op(e3, e2))) | spl140_249), inference(avatar_component_clause, [], [f2606])).
fof(f5986, plain, (spl140_90 | ~ spl140_18 | ~ spl140_256), inference(avatar_split_clause, [], [f5985, f2637, f1309, f1611])).
fof(f2637, plain, (spl140_256 <=> (e4 = op(e1, op(e4, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl140_256])])).
fof(f5985, plain, ((e4 = op(e1, e2)) | (~ spl140_18 | ~ spl140_256)), inference(forward_demodulation, [], [f2638, f1311])).
fof(f1311, plain, ((e2 = op(e4, e1)) | ~ spl140_18), inference(avatar_component_clause, [], [f1309])).
fof(f2638, plain, ((e4 = op(e1, op(e4, e1))) | ~ spl140_256), inference(avatar_component_clause, [], [f2637])).
fof(f5984, plain, (~ spl140_1 | ~ spl140_25 | spl140_264), inference(avatar_split_clause, [], [f5983, f2673, f1338, f1238])).
fof(f1338, plain, (spl140_25 <=> (e4 = op(e4, unit))), introduced(avatar_definition, [new_symbols(naming, [spl140_25])])).
fof(f2673, plain, (spl140_264 <=> (unit = op(e4, op(e4, unit)))), introduced(avatar_definition, [new_symbols(naming, [spl140_264])])).
fof(f5983, plain, (~ (op(e4, e4) = unit) | (~ spl140_25 | spl140_264)), inference(forward_demodulation, [], [f2674, f1340])).
fof(f1340, plain, ((e4 = op(e4, unit)) | ~ spl140_25), inference(avatar_component_clause, [], [f1338])).
fof(f2674, plain, (~ (unit = op(e4, op(e4, unit))) | spl140_264), inference(avatar_component_clause, [], [f2673])).
fof(f5982, plain, (spl140_31 | ~ spl140_49 | ~ spl140_266), inference(avatar_split_clause, [], [f5981, f2682, f1439, f1364])).
fof(f1439, plain, (spl140_49 <=> (e3 = op(e3, unit))), introduced(avatar_definition, [new_symbols(naming, [spl140_49])])).
fof(f2682, plain, (spl140_266 <=> (unit = op(e3, op(e3, unit)))), introduced(avatar_definition, [new_symbols(naming, [spl140_266])])).
fof(f5981, plain, ((op(e3, e3) = unit) | (~ spl140_49 | ~ spl140_266)), inference(forward_demodulation, [], [f2684, f1441])).
fof(f1441, plain, ((e3 = op(e3, unit)) | ~ spl140_49), inference(avatar_component_clause, [], [f1439])).
fof(f2684, plain, ((unit = op(e3, op(e3, unit))) | ~ spl140_266), inference(avatar_component_clause, [], [f2682])).
fof(f5972, plain, (~ spl140_14 | ~ spl140_28 | spl140_349), inference(avatar_split_clause, [], [f5971, f3182, f1351, f1292])).
fof(f3182, plain, (spl140_349 <=> (e3 = op(e4, op(e3, e4)))), introduced(avatar_definition, [new_symbols(naming, [spl140_349])])).
fof(f5971, plain, (~ (e3 = op(e4, e2)) | (~ spl140_28 | spl140_349)), inference(forward_demodulation, [], [f3184, f1353])).
fof(f1353, plain, ((e2 = op(e3, e4)) | ~ spl140_28), inference(avatar_component_clause, [], [f1351])).
fof(f3184, plain, (~ (e3 = op(e4, op(e3, e4))) | spl140_349), inference(avatar_component_clause, [], [f3182])).
fof(f5961, plain, (~ spl140_70 | ~ spl140_52 | spl140_350), inference(avatar_split_clause, [], [f5761, f3187, f1452, f1527])).
fof(f5761, plain, (~ (e4 = op(e2, e1)) | (~ spl140_52 | spl140_350)), inference(forward_demodulation, [], [f3188, f1454])).
fof(f5955, plain, (~ spl140_83 | spl140_89 | ~ spl140_241), inference(avatar_contradiction_clause, [], [f5954])).
fof(f5954, plain, ($false | (~ spl140_83 | spl140_89 | ~ spl140_241)), inference(subsumption_resolution, [], [f5953, f1608])).
fof(f1608, plain, (~ (e3 = op(e1, e2)) | spl140_89), inference(avatar_component_clause, [], [f1607])).
fof(f1607, plain, (spl140_89 <=> (e3 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl140_89])])).
fof(f5953, plain, ((e3 = op(e1, e2)) | (~ spl140_83 | ~ spl140_241)), inference(forward_demodulation, [], [f2573, f1584])).
fof(f2573, plain, ((e3 = op(e1, op(e1, e3))) | ~ spl140_241), inference(avatar_component_clause, [], [f2571])).
fof(f2571, plain, (spl140_241 <=> (e3 = op(e1, op(e1, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl140_241])])).
fof(f5936, plain, (~ spl140_4 | ~ spl140_14), inference(avatar_split_clause, [], [f5931, f1292, f1250])).
fof(f5931, plain, (~ (e3 = op(e4, e4)) | ~ spl140_14), inference(backward_demodulation, [], [f476, f1294])).
fof(f5929, plain, (spl140_14 | ~ spl140_28 | ~ spl140_349), inference(avatar_split_clause, [], [f5926, f3182, f1351, f1292])).
fof(f5926, plain, ((e3 = op(e4, e2)) | (~ spl140_28 | ~ spl140_349)), inference(backward_demodulation, [], [f3183, f1353])).
fof(f3183, plain, ((e3 = op(e4, op(e3, e4))) | ~ spl140_349), inference(avatar_component_clause, [], [f3182])).
fof(f5913, plain, (~ spl140_26 | ~ spl140_41), inference(avatar_split_clause, [], [f5908, f1406, f1343])).
fof(f5908, plain, (~ (op(e3, e4) = unit) | ~ spl140_41), inference(backward_demodulation, [], [f464, f1408])).
fof(f1408, plain, ((op(e3, e1) = unit) | ~ spl140_41), inference(avatar_component_clause, [], [f1406])).
fof(f5897, plain, (~ spl140_45 | ~ spl140_79 | spl140_332), inference(avatar_split_clause, [], [f5891, f3025, f1565, f1422])).
fof(f5891, plain, (~ (e4 = op(e3, e1)) | (~ spl140_79 | spl140_332)), inference(backward_demodulation, [], [f3026, f1567])).
fof(f3026, plain, (~ (e4 = op(op(e1, e4), e1)) | spl140_332), inference(avatar_component_clause, [], [f3025])).
fof(f5880, plain, (~ spl140_78 | ~ spl140_83), inference(avatar_split_clause, [], [f5871, f1582, f1561])).
fof(f5871, plain, (~ (e2 = op(e1, e4)) | ~ spl140_83), inference(backward_demodulation, [], [f447, f1584])).
fof(f447, plain, ~ (op(e1, e3) = op(e1, e4)), inference(cnf_transformation, [], [f5])).
fof(f5879, plain, (~ spl140_33 | ~ spl140_83), inference(avatar_split_clause, [], [f5870, f1582, f1372])).
fof(f1372, plain, (spl140_33 <=> (e2 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl140_33])])).
fof(f5870, plain, (~ (e2 = op(e3, e3)) | ~ spl140_83), inference(backward_demodulation, [], [f413, f1584])).
fof(f5867, plain, (~ spl140_81 | ~ spl140_86), inference(avatar_split_clause, [], [f5861, f1595, f1574])).
fof(f5861, plain, (~ (op(e1, e3) = unit) | ~ spl140_86), inference(backward_demodulation, [], [f445, f1597])).
fof(f445, plain, ~ (op(e1, e2) = op(e1, e3)), inference(cnf_transformation, [], [f5])).
fof(f5854, plain, (spl140_1 | ~ spl140_25 | ~ spl140_264), inference(avatar_split_clause, [], [f5853, f2673, f1338, f1238])).
fof(f5853, plain, ((op(e4, e4) = unit) | (~ spl140_25 | ~ spl140_264)), inference(forward_demodulation, [], [f2675, f1340])).
fof(f2675, plain, ((unit = op(e4, op(e4, unit))) | ~ spl140_264), inference(avatar_component_clause, [], [f2673])).
fof(f5843, plain, (~ spl140_353 | spl140_255 | ~ spl140_304), inference(avatar_split_clause, [], [f5578, f2903, f2633, f3200])).
fof(f2633, plain, (spl140_255 <=> (e1 = op(e4, op(e4, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl140_255])])).
fof(f2903, plain, (spl140_304 <=> (op(e1, e4) = op(e4, e1))), introduced(avatar_definition, [new_symbols(naming, [spl140_304])])).
fof(f5578, plain, (~ (e1 = op(e4, op(e1, e4))) | (spl140_255 | ~ spl140_304)), inference(backward_demodulation, [], [f2634, f2905])).
fof(f2905, plain, ((op(e1, e4) = op(e4, e1)) | ~ spl140_304), inference(avatar_component_clause, [], [f2903])).
fof(f2634, plain, (~ (e1 = op(e4, op(e4, e1))) | spl140_255), inference(avatar_component_clause, [], [f2633])).
fof(f5835, plain, (spl140_78 | ~ spl140_18 | ~ spl140_304), inference(avatar_split_clause, [], [f5834, f2903, f1309, f1561])).
fof(f5834, plain, ((e2 = op(e1, e4)) | (~ spl140_18 | ~ spl140_304)), inference(forward_demodulation, [], [f2905, f1311])).
fof(f5833, plain, (~ spl140_79 | ~ spl140_69 | ~ spl140_304), inference(avatar_split_clause, [], [f5832, f2903, f1523, f1565])).
fof(f5832, plain, (~ (e3 = op(e1, e4)) | (~ spl140_69 | ~ spl140_304)), inference(forward_demodulation, [], [f5576, f1525])).
fof(f5576, plain, (~ (op(e1, e4) = op(e2, e1)) | ~ spl140_304), inference(backward_demodulation, [], [f396, f2905])).
fof(f396, plain, ~ (op(e2, e1) = op(e4, e1)), inference(cnf_transformation, [], [f5])).
fof(f5804, plain, (~ spl140_56 | ~ spl140_81), inference(avatar_split_clause, [], [f5799, f1574, f1469])).
fof(f1469, plain, (spl140_56 <=> (op(e2, e3) = unit)), introduced(avatar_definition, [new_symbols(naming, [spl140_56])])).
fof(f5799, plain, (~ (op(e2, e3) = unit) | ~ spl140_81), inference(backward_demodulation, [], [f412, f1576])).
fof(f412, plain, ~ (op(e1, e3) = op(e2, e3)), inference(cnf_transformation, [], [f5])).
fof(f5797, plain, (~ spl140_85 | ~ spl140_95), inference(avatar_split_clause, [], [f5796, f1632, f1590])).
fof(f5796, plain, (~ (e4 = op(e1, e3)) | ~ spl140_95), inference(backward_demodulation, [], [f443, f1634])).
fof(f5791, plain, (~ spl140_37 | spl140_43 | ~ spl140_248), inference(avatar_contradiction_clause, [], [f5790])).
fof(f5790, plain, ($false | (~ spl140_37 | spl140_43 | ~ spl140_248)), inference(subsumption_resolution, [], [f5789, f1415])).
fof(f1415, plain, (~ (e2 = op(e3, e1)) | spl140_43), inference(avatar_component_clause, [], [f1414])).
fof(f5789, plain, ((e2 = op(e3, e1)) | (~ spl140_37 | ~ spl140_248)), inference(forward_demodulation, [], [f2604, f1391])).
fof(f2604, plain, ((e2 = op(e3, op(e3, e2))) | ~ spl140_248), inference(avatar_component_clause, [], [f2602])).
fof(f2602, plain, (spl140_248 <=> (e2 = op(e3, op(e3, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl140_248])])).
fof(f5751, plain, (~ spl140_60 | ~ spl140_65), inference(avatar_split_clause, [], [f5495, f1506, f1485])).
fof(f5495, plain, (~ (e4 = op(e2, e3)) | ~ spl140_65), inference(backward_demodulation, [], [f455, f1508])).
fof(f1508, plain, ((e4 = op(e2, e2)) | ~ spl140_65), inference(avatar_component_clause, [], [f1506])).
fof(f455, plain, ~ (op(e2, e2) = op(e2, e3)), inference(cnf_transformation, [], [f5])).
fof(f5749, plain, (~ spl140_26 | ~ spl140_49 | ~ spl140_348), inference(avatar_contradiction_clause, [], [f5748])).
fof(f5748, plain, ($false | (~ spl140_26 | ~ spl140_49 | ~ spl140_348)), inference(subsumption_resolution, [], [f5747, f487])).
fof(f487, plain, ~ (e3 = e4), inference(cnf_transformation, [], [f6])).
fof(f5747, plain, ((e3 = e4) | (~ spl140_26 | ~ spl140_49 | ~ spl140_348)), inference(forward_demodulation, [], [f5740, f1441])).
fof(f5740, plain, ((e4 = op(e3, unit)) | (~ spl140_26 | ~ spl140_348)), inference(backward_demodulation, [], [f3180, f1345])).
fof(f1345, plain, ((op(e3, e4) = unit) | ~ spl140_26), inference(avatar_component_clause, [], [f1343])).
fof(f3180, plain, ((e4 = op(e3, op(e3, e4))) | ~ spl140_348), inference(avatar_component_clause, [], [f3178])).
fof(f3178, plain, (spl140_348 <=> (e4 = op(e3, op(e3, e4)))), introduced(avatar_definition, [new_symbols(naming, [spl140_348])])).
fof(f5720, plain, (~ spl140_56 | ~ spl140_113 | ~ spl140_324), inference(avatar_contradiction_clause, [], [f5719])).
fof(f5719, plain, ($false | (~ spl140_56 | ~ spl140_113 | ~ spl140_324)), inference(subsumption_resolution, [], [f5718, f485])).
fof(f485, plain, ~ (e2 = e3), inference(cnf_transformation, [], [f6])).
fof(f5718, plain, ((e2 = e3) | (~ spl140_56 | ~ spl140_113 | ~ spl140_324)), inference(forward_demodulation, [], [f5713, f1710])).
fof(f1710, plain, ((e2 = op(unit, e2)) | ~ spl140_113), inference(avatar_component_clause, [], [f1708])).
fof(f1708, plain, (spl140_113 <=> (e2 = op(unit, e2))), introduced(avatar_definition, [new_symbols(naming, [spl140_113])])).
fof(f5713, plain, ((e3 = op(unit, e2)) | (~ spl140_56 | ~ spl140_324)), inference(backward_demodulation, [], [f2992, f1471])).
fof(f1471, plain, ((op(e2, e3) = unit) | ~ spl140_56), inference(avatar_component_clause, [], [f1469])).
fof(f5714, plain, (~ spl140_51 | ~ spl140_56), inference(avatar_split_clause, [], [f5709, f1469, f1448])).
fof(f5709, plain, (~ (op(e2, e4) = unit) | ~ spl140_56), inference(backward_demodulation, [], [f457, f1471])).
fof(f457, plain, ~ (op(e2, e3) = op(e2, e4)), inference(cnf_transformation, [], [f5])).
fof(f5704, plain, (~ spl140_57 | ~ spl140_89 | spl140_251), inference(avatar_split_clause, [], [f5703, f2615, f1607, f1473])).
fof(f1473, plain, (spl140_57 <=> (e1 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl140_57])])).
fof(f5703, plain, (~ (e1 = op(e2, e3)) | (~ spl140_89 | spl140_251)), inference(forward_demodulation, [], [f2617, f1609])).
fof(f1609, plain, ((e3 = op(e1, e2)) | ~ spl140_89), inference(avatar_component_clause, [], [f1607])).
fof(f5690, plain, (~ spl140_57 | ~ spl140_37 | spl140_312), inference(avatar_split_clause, [], [f5689, f2938, f1389, f1473])).
fof(f5689, plain, (~ (e1 = op(e2, e3)) | (~ spl140_37 | spl140_312)), inference(forward_demodulation, [], [f2939, f1391])).
fof(f2939, plain, (~ (op(e2, e3) = op(e3, e2)) | spl140_312), inference(avatar_component_clause, [], [f2938])).
fof(f5661, plain, (~ spl140_57 | ~ spl140_7), inference(avatar_split_clause, [], [f5660, f1263, f1473])).
fof(f5660, plain, (~ (e1 = op(e2, e3)) | ~ spl140_7), inference(forward_demodulation, [], [f416, f1265])).
fof(f416, plain, ~ (op(e2, e3) = op(e4, e3)), inference(cnf_transformation, [], [f5])).
fof(f5644, plain, (~ spl140_19 | ~ spl140_27 | spl140_349), inference(avatar_split_clause, [], [f5638, f3182, f1347, f1313])).
fof(f5638, plain, (~ (e3 = op(e4, e1)) | (~ spl140_27 | spl140_349)), inference(backward_demodulation, [], [f3184, f1349])).
fof(f1349, plain, ((e1 = op(e3, e4)) | ~ spl140_27), inference(avatar_component_clause, [], [f1347])).
fof(f5640, plain, (spl140_7 | ~ spl140_27 | ~ spl140_298), inference(avatar_split_clause, [], [f5634, f2877, f1347, f1263])).
fof(f2877, plain, (spl140_298 <=> (op(e3, e4) = op(e4, e3))), introduced(avatar_definition, [new_symbols(naming, [spl140_298])])).
fof(f5634, plain, ((e1 = op(e4, e3)) | (~ spl140_27 | ~ spl140_298)), inference(backward_demodulation, [], [f2879, f1349])).
fof(f2879, plain, ((op(e3, e4) = op(e4, e3)) | ~ spl140_298), inference(avatar_component_clause, [], [f2877])).
fof(f5639, plain, (~ spl140_2 | ~ spl140_27), inference(avatar_split_clause, [], [f5629, f1347, f1242])).
fof(f1242, plain, (spl140_2 <=> (e1 = op(e4, e4))), introduced(avatar_definition, [new_symbols(naming, [spl140_2])])).
fof(f5629, plain, (~ (e1 = op(e4, e4)) | ~ spl140_27), inference(backward_demodulation, [], [f427, f1349])).
fof(f427, plain, ~ (op(e3, e4) = op(e4, e4)), inference(cnf_transformation, [], [f5])).
fof(f5628, plain, (~ spl140_31 | ~ spl140_33), inference(avatar_contradiction_clause, [], [f5627])).
fof(f5627, plain, ($false | (~ spl140_31 | ~ spl140_33)), inference(subsumption_resolution, [], [f5624, f1032])).
fof(f1032, plain, ~ (e2 = unit), inference(definition_unfolding, [], [f479, f377])).
fof(f377, plain, (e0 = unit), inference(cnf_transformation, [], [f4])).
fof(f4, plain, (e0 = unit), file('/home/ubuntu/library/tptp/Problems/ALG/ALG047+1.p', ax4)).
fof(f479, plain, ~ (e0 = e2), inference(cnf_transformation, [], [f6])).
fof(f5624, plain, ((e2 = unit) | (~ spl140_31 | ~ spl140_33)), inference(backward_demodulation, [], [f1366, f1374])).
fof(f1374, plain, ((e2 = op(e3, e3)) | ~ spl140_33), inference(avatar_component_clause, [], [f1372])).
fof(f5616, plain, (~ spl140_54 | ~ spl140_69), inference(avatar_split_clause, [], [f5613, f1523, f1460])).
fof(f1460, plain, (spl140_54 <=> (e3 = op(e2, e4))), introduced(avatar_definition, [new_symbols(naming, [spl140_54])])).
fof(f5613, plain, (~ (e3 = op(e2, e4)) | ~ spl140_69), inference(backward_demodulation, [], [f454, f1525])).
fof(f454, plain, ~ (op(e2, e1) = op(e2, e4)), inference(cnf_transformation, [], [f5])).
fof(f5610, plain, (spl140_18 | ~ spl140_78 | ~ spl140_304), inference(avatar_split_clause, [], [f5603, f2903, f1561, f1309])).
fof(f5603, plain, ((e2 = op(e4, e1)) | (~ spl140_78 | ~ spl140_304)), inference(backward_demodulation, [], [f2905, f1563])).
fof(f1563, plain, ((e2 = op(e1, e4)) | ~ spl140_78), inference(avatar_component_clause, [], [f1561])).
fof(f5609, plain, (~ spl140_28 | ~ spl140_78), inference(avatar_split_clause, [], [f5599, f1561, f1351])).
fof(f5599, plain, (~ (e2 = op(e3, e4)) | ~ spl140_78), inference(backward_demodulation, [], [f423, f1563])).
fof(f423, plain, ~ (op(e1, e4) = op(e3, e4)), inference(cnf_transformation, [], [f5])).
fof(f5597, plain, (~ spl140_76 | ~ spl140_91), inference(avatar_split_clause, [], [f5595, f1616, f1553])).
fof(f1616, plain, (spl140_91 <=> (op(e1, e1) = unit)), introduced(avatar_definition, [new_symbols(naming, [spl140_91])])).
fof(f5595, plain, (~ (op(e1, e4) = unit) | ~ spl140_91), inference(backward_demodulation, [], [f444, f1618])).
fof(f1618, plain, ((op(e1, e1) = unit) | ~ spl140_91), inference(avatar_component_clause, [], [f1616])).
fof(f5596, plain, (~ spl140_66 | ~ spl140_91), inference(avatar_split_clause, [], [f5593, f1616, f1511])).
fof(f5593, plain, (~ (op(e2, e1) = unit) | ~ spl140_91), inference(backward_demodulation, [], [f392, f1618])).
fof(f5592, plain, (spl140_69 | ~ spl140_57 | ~ spl140_239), inference(avatar_split_clause, [], [f5591, f2562, f1473, f1523])).
fof(f2562, plain, (spl140_239 <=> (e3 = op(e2, op(e2, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl140_239])])).
fof(f5591, plain, ((e3 = op(e2, e1)) | (~ spl140_57 | ~ spl140_239)), inference(forward_demodulation, [], [f2564, f1475])).
fof(f1475, plain, ((e1 = op(e2, e3)) | ~ spl140_57), inference(avatar_component_clause, [], [f1473])).
fof(f2564, plain, ((e3 = op(e2, op(e2, e3))) | ~ spl140_239), inference(avatar_component_clause, [], [f2562])).
fof(f5585, plain, (spl140_91 | ~ spl140_97 | ~ spl140_270), inference(avatar_split_clause, [], [f5584, f2700, f1641, f1616])).
fof(f1641, plain, (spl140_97 <=> (e1 = op(e1, unit))), introduced(avatar_definition, [new_symbols(naming, [spl140_97])])).
fof(f2700, plain, (spl140_270 <=> (unit = op(e1, op(e1, unit)))), introduced(avatar_definition, [new_symbols(naming, [spl140_270])])).
fof(f5584, plain, ((op(e1, e1) = unit) | (~ spl140_97 | ~ spl140_270)), inference(forward_demodulation, [], [f2702, f1643])).
fof(f1643, plain, ((e1 = op(e1, unit)) | ~ spl140_97), inference(avatar_component_clause, [], [f1641])).
fof(f2702, plain, ((unit = op(e1, op(e1, unit))) | ~ spl140_270), inference(avatar_component_clause, [], [f2700])).
fof(f5571, plain, (~ spl140_19 | ~ spl140_85 | spl140_334), inference(avatar_split_clause, [], [f5570, f3034, f1590, f1313])).
fof(f5570, plain, (~ (e3 = op(e4, e1)) | (~ spl140_85 | spl140_334)), inference(forward_demodulation, [], [f3035, f1592])).
fof(f1592, plain, ((e4 = op(e1, e3)) | ~ spl140_85), inference(avatar_component_clause, [], [f1590])).
fof(f3035, plain, (~ (e3 = op(op(e1, e3), e1)) | spl140_334), inference(avatar_component_clause, [], [f3034])).
fof(f5558, plain, (~ spl140_27 | ~ spl140_57 | ~ spl140_312), inference(avatar_split_clause, [], [f5509, f2938, f1473, f1347])).
fof(f5509, plain, (~ (e1 = op(e3, e4)) | (~ spl140_57 | ~ spl140_312)), inference(backward_demodulation, [], [f5446, f1475])).
fof(f5446, plain, (~ (op(e2, e3) = op(e3, e4)) | ~ spl140_312), inference(backward_demodulation, [], [f466, f2940])).
fof(f5550, plain, (~ spl140_6 | ~ spl140_31), inference(avatar_split_clause, [], [f5549, f1364, f1259])).
fof(f1259, plain, (spl140_6 <=> (op(e4, e3) = unit)), introduced(avatar_definition, [new_symbols(naming, [spl140_6])])).
fof(f5549, plain, (~ (op(e4, e3) = unit) | ~ spl140_31), inference(forward_demodulation, [], [f417, f1366])).
fof(f417, plain, ~ (op(e3, e3) = op(e4, e3)), inference(cnf_transformation, [], [f5])).
fof(f5546, plain, (~ spl140_26 | ~ spl140_11 | ~ spl140_298), inference(avatar_split_clause, [], [f5545, f2877, f1280, f1343])).
fof(f5545, plain, (~ (op(e3, e4) = unit) | (~ spl140_11 | ~ spl140_298)), inference(forward_demodulation, [], [f5544, f1282])).
fof(f5544, plain, (~ (op(e3, e4) = op(e4, e2)) | ~ spl140_298), inference(forward_demodulation, [], [f475, f2879])).
fof(f475, plain, ~ (op(e4, e2) = op(e4, e3)), inference(cnf_transformation, [], [f5])).
fof(f5511, plain, (~ spl140_43 | ~ spl140_57 | spl140_240), inference(avatar_split_clause, [], [f5504, f2566, f1473, f1414])).
fof(f5504, plain, (~ (e2 = op(e3, e1)) | (~ spl140_57 | spl140_240)), inference(backward_demodulation, [], [f2568, f1475])).
fof(f5510, plain, (~ spl140_32 | ~ spl140_57), inference(avatar_split_clause, [], [f5499, f1473, f1368])).
fof(f1368, plain, (spl140_32 <=> (e1 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl140_32])])).
fof(f5499, plain, (~ (e1 = op(e3, e3)) | ~ spl140_57), inference(backward_demodulation, [], [f415, f1475])).
fof(f5498, plain, (~ spl140_40 | ~ spl140_65), inference(avatar_split_clause, [], [f5497, f1506, f1401])).
fof(f5497, plain, (~ (e4 = op(e3, e2)) | ~ spl140_65), inference(backward_demodulation, [], [f405, f1508])).
fof(f5482, plain, (~ spl140_26 | ~ spl140_76), inference(avatar_split_clause, [], [f5478, f1553, f1343])).
fof(f5478, plain, (~ (op(e3, e4) = unit) | ~ spl140_76), inference(backward_demodulation, [], [f423, f1555])).
fof(f5476, plain, (~ spl140_78 | ~ spl140_93), inference(avatar_split_clause, [], [f5472, f1624, f1561])).
fof(f5472, plain, (~ (e2 = op(e1, e4)) | ~ spl140_93), inference(backward_demodulation, [], [f444, f1626])).
fof(f5468, plain, (~ spl140_40 | ~ spl140_8 | spl140_238), inference(avatar_split_clause, [], [f5467, f2557, f1267, f1401])).
fof(f1267, plain, (spl140_8 <=> (e2 = op(e4, e3))), introduced(avatar_definition, [new_symbols(naming, [spl140_8])])).
fof(f5467, plain, (~ (e4 = op(e3, e2)) | (~ spl140_8 | spl140_238)), inference(forward_demodulation, [], [f2559, f1269])).
fof(f1269, plain, ((e2 = op(e4, e3)) | ~ spl140_8), inference(avatar_component_clause, [], [f1267])).
fof(f5452, plain, (~ spl140_19 | spl140_27 | ~ spl140_306), inference(avatar_contradiction_clause, [], [f5451])).
fof(f5451, plain, ($false | (~ spl140_19 | spl140_27 | ~ spl140_306)), inference(subsumption_resolution, [], [f5450, f1348])).
fof(f1348, plain, (~ (e1 = op(e3, e4)) | spl140_27), inference(avatar_component_clause, [], [f1347])).
fof(f5450, plain, ((e1 = op(e3, e4)) | (~ spl140_19 | ~ spl140_306)), inference(forward_demodulation, [], [f2913, f1315])).
fof(f1315, plain, ((e3 = op(e4, e1)) | ~ spl140_19), inference(avatar_component_clause, [], [f1313])).
fof(f2913, plain, ((e1 = op(op(e4, e1), e4)) | ~ spl140_306), inference(avatar_component_clause, [], [f2911])).
fof(f2911, plain, (spl140_306 <=> (e1 = op(op(e4, e1), e4))), introduced(avatar_definition, [new_symbols(naming, [spl140_306])])).
fof(f5443, plain, (spl140_45 | ~ spl140_85 | ~ spl140_315), inference(avatar_split_clause, [], [f5442, f2951, f1590, f1422])).
fof(f5442, plain, ((e4 = op(e3, e1)) | (~ spl140_85 | ~ spl140_315)), inference(forward_demodulation, [], [f2953, f1592])).
fof(f5434, plain, (~ spl140_79 | ~ spl140_54), inference(avatar_split_clause, [], [f5433, f1460, f1565])).
fof(f5433, plain, (~ (e3 = op(e1, e4)) | ~ spl140_54), inference(forward_demodulation, [], [f422, f1462])).
fof(f1462, plain, ((e3 = op(e2, e4)) | ~ spl140_54), inference(avatar_component_clause, [], [f1460])).
fof(f422, plain, ~ (op(e1, e4) = op(e2, e4)), inference(cnf_transformation, [], [f5])).
fof(f5430, plain, (~ spl140_79 | ~ spl140_89), inference(avatar_split_clause, [], [f5429, f1607, f1565])).
fof(f5429, plain, (~ (e3 = op(e1, e4)) | ~ spl140_89), inference(forward_demodulation, [], [f446, f1609])).
fof(f5428, plain, (~ spl140_79 | ~ spl140_19 | spl140_304), inference(avatar_split_clause, [], [f5427, f2903, f1313, f1565])).
fof(f5427, plain, (~ (e3 = op(e1, e4)) | (~ spl140_19 | spl140_304)), inference(forward_demodulation, [], [f2904, f1315])).
fof(f2904, plain, (~ (op(e1, e4) = op(e4, e1)) | spl140_304), inference(avatar_component_clause, [], [f2903])).
fof(f5418, plain, (~ spl140_36 | ~ spl140_11), inference(avatar_split_clause, [], [f5386, f1280, f1385])).
fof(f5386, plain, (~ (op(e3, e2) = unit) | ~ spl140_11), inference(forward_demodulation, [], [f407, f1282])).
fof(f407, plain, ~ (op(e3, e2) = op(e4, e2)), inference(cnf_transformation, [], [f5])).
fof(f5417, plain, (~ spl140_37 | ~ spl140_89 | spl140_335), inference(avatar_split_clause, [], [f5068, f3039, f1607, f1389])).
fof(f3039, plain, (spl140_335 <=> (e1 = op(op(e1, e2), e2))), introduced(avatar_definition, [new_symbols(naming, [spl140_335])])).
fof(f5068, plain, (~ (e1 = op(e3, e2)) | (~ spl140_89 | spl140_335)), inference(backward_demodulation, [], [f3041, f1609])).
fof(f3041, plain, (~ (e1 = op(op(e1, e2), e2)) | spl140_335), inference(avatar_component_clause, [], [f3039])).
fof(f5416, plain, (~ spl140_35 | ~ spl140_85), inference(avatar_split_clause, [], [f5415, f1590, f1380])).
fof(f5415, plain, (~ (e4 = op(e3, e3)) | ~ spl140_85), inference(forward_demodulation, [], [f413, f1592])).
fof(f5411, plain, (~ spl140_8 | spl140_14 | ~ spl140_237), inference(avatar_contradiction_clause, [], [f5410])).
fof(f5410, plain, ($false | (~ spl140_8 | spl140_14 | ~ spl140_237)), inference(subsumption_resolution, [], [f5409, f1293])).
fof(f1293, plain, (~ (e3 = op(e4, e2)) | spl140_14), inference(avatar_component_clause, [], [f1292])).
fof(f5409, plain, ((e3 = op(e4, e2)) | (~ spl140_8 | ~ spl140_237)), inference(forward_demodulation, [], [f2555, f1269])).
fof(f5397, plain, (~ spl140_54 | spl140_60 | ~ spl140_350), inference(avatar_contradiction_clause, [], [f5396])).
fof(f5396, plain, ($false | (~ spl140_54 | spl140_60 | ~ spl140_350)), inference(subsumption_resolution, [], [f5395, f1486])).
fof(f1486, plain, (~ (e4 = op(e2, e3)) | spl140_60), inference(avatar_component_clause, [], [f1485])).
fof(f5395, plain, ((e4 = op(e2, e3)) | (~ spl140_54 | ~ spl140_350)), inference(forward_demodulation, [], [f3189, f1462])).
fof(f5370, plain, (~ spl140_4 | ~ spl140_19), inference(avatar_split_clause, [], [f5365, f1313, f1250])).
fof(f5365, plain, (~ (e3 = op(e4, e4)) | ~ spl140_19), inference(backward_demodulation, [], [f474, f1315])).
fof(f474, plain, ~ (op(e4, e1) = op(e4, e4)), inference(cnf_transformation, [], [f5])).
fof(f5341, plain, (~ spl140_12 | ~ spl140_62), inference(avatar_split_clause, [], [f5335, f1494, f1284])).
fof(f1284, plain, (spl140_12 <=> (e1 = op(e4, e2))), introduced(avatar_definition, [new_symbols(naming, [spl140_12])])).
fof(f5335, plain, (~ (e1 = op(e4, e2)) | ~ spl140_62), inference(backward_demodulation, [], [f406, f1496])).
fof(f5332, plain, (~ spl140_60 | ~ spl140_70), inference(avatar_split_clause, [], [f5325, f1527, f1485])).
fof(f5325, plain, (~ (e4 = op(e2, e3)) | ~ spl140_70), inference(backward_demodulation, [], [f453, f1529])).
fof(f1529, plain, ((e4 = op(e2, e1)) | ~ spl140_70), inference(avatar_component_clause, [], [f1527])).
fof(f453, plain, ~ (op(e2, e1) = op(e2, e3)), inference(cnf_transformation, [], [f5])).
fof(f5321, plain, (spl140_19 | ~ spl140_85 | ~ spl140_334), inference(avatar_split_clause, [], [f5318, f3034, f1590, f1313])).
fof(f5318, plain, ((e3 = op(e4, e1)) | (~ spl140_85 | ~ spl140_334)), inference(backward_demodulation, [], [f3036, f1592])).
fof(f5319, plain, (~ spl140_60 | ~ spl140_85), inference(avatar_split_clause, [], [f5314, f1590, f1485])).
fof(f5314, plain, (~ (e4 = op(e2, e3)) | ~ spl140_85), inference(backward_demodulation, [], [f412, f1592])).
fof(f5313, plain, (~ spl140_81 | ~ spl140_91), inference(avatar_split_clause, [], [f5311, f1616, f1574])).
fof(f5311, plain, (~ (op(e1, e3) = unit) | ~ spl140_91), inference(backward_demodulation, [], [f443, f1618])).
fof(f5312, plain, (~ spl140_16 | ~ spl140_91), inference(avatar_split_clause, [], [f5310, f1616, f1301])).
fof(f1301, plain, (spl140_16 <=> (op(e4, e1) = unit)), introduced(avatar_definition, [new_symbols(naming, [spl140_16])])).
fof(f5310, plain, (~ (op(e4, e1) = unit) | ~ spl140_91), inference(backward_demodulation, [], [f394, f1618])).
fof(f394, plain, ~ (op(e1, e1) = op(e4, e1)), inference(cnf_transformation, [], [f5])).
fof(f5308, plain, (spl140_28 | ~ spl140_40 | ~ spl140_248), inference(avatar_contradiction_clause, [], [f5307])).
fof(f5307, plain, ($false | (spl140_28 | ~ spl140_40 | ~ spl140_248)), inference(subsumption_resolution, [], [f5306, f1352])).
fof(f5306, plain, ((e2 = op(e3, e4)) | (~ spl140_40 | ~ spl140_248)), inference(forward_demodulation, [], [f2604, f1403])).
fof(f1403, plain, ((e4 = op(e3, e2)) | ~ spl140_40), inference(avatar_component_clause, [], [f1401])).
fof(f5305, plain, (spl140_54 | ~ spl140_40 | ~ spl140_249), inference(avatar_split_clause, [], [f5304, f2606, f1401, f1460])).
fof(f5304, plain, ((e3 = op(e2, e4)) | (~ spl140_40 | ~ spl140_249)), inference(forward_demodulation, [], [f2607, f1403])).
fof(f2607, plain, ((e3 = op(e2, op(e3, e2))) | ~ spl140_249), inference(avatar_component_clause, [], [f2606])).
fof(f5291, plain, (~ spl140_69 | ~ spl140_89 | spl140_325), inference(avatar_split_clause, [], [f5290, f2995, f1607, f1523])).
fof(f2995, plain, (spl140_325 <=> (op(e1, e2) = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl140_325])])).
fof(f5290, plain, (~ (e3 = op(e2, e1)) | (~ spl140_89 | spl140_325)), inference(forward_demodulation, [], [f2996, f1609])).
fof(f2996, plain, (~ (op(e1, e2) = op(e2, e1)) | spl140_325), inference(avatar_component_clause, [], [f2995])).
fof(f5289, plain, (~ spl140_52 | ~ spl140_78 | spl140_331), inference(avatar_split_clause, [], [f5288, f3021, f1561, f1452])).
fof(f5288, plain, (~ (e1 = op(e2, e4)) | (~ spl140_78 | spl140_331)), inference(forward_demodulation, [], [f3023, f1563])).
fof(f5283, plain, (~ spl140_12 | ~ spl140_78 | spl140_353), inference(avatar_split_clause, [], [f5282, f3200, f1561, f1284])).
fof(f5282, plain, (~ (e1 = op(e4, e2)) | (~ spl140_78 | spl140_353)), inference(forward_demodulation, [], [f3202, f1563])).
fof(f3202, plain, (~ (e1 = op(e4, op(e1, e4))) | spl140_353), inference(avatar_component_clause, [], [f3200])).
fof(f5278, plain, (spl140_70 | ~ spl140_78 | ~ spl140_332), inference(avatar_split_clause, [], [f5236, f3025, f1561, f1527])).
fof(f5236, plain, ((e4 = op(e2, e1)) | (~ spl140_78 | ~ spl140_332)), inference(forward_demodulation, [], [f3027, f1563])).
fof(f5275, plain, (spl140_60 | ~ spl140_40 | ~ spl140_312), inference(avatar_split_clause, [], [f5274, f2938, f1401, f1485])).
fof(f5274, plain, ((e4 = op(e2, e3)) | (~ spl140_40 | ~ spl140_312)), inference(forward_demodulation, [], [f2940, f1403])).
fof(f5259, plain, (~ spl140_14 | ~ spl140_89), inference(avatar_split_clause, [], [f5258, f1607, f1292])).
fof(f5258, plain, (~ (e3 = op(e4, e2)) | ~ spl140_89), inference(forward_demodulation, [], [f404, f1609])).
fof(f5257, plain, (~ spl140_351 | spl140_246 | ~ spl140_301), inference(avatar_split_clause, [], [f5256, f2890, f2593, f3191])).
fof(f3191, plain, (spl140_351 <=> (e2 = op(e4, op(e2, e4)))), introduced(avatar_definition, [new_symbols(naming, [spl140_351])])).
fof(f2593, plain, (spl140_246 <=> (e2 = op(e4, op(e4, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl140_246])])).
fof(f5256, plain, (~ (e2 = op(e4, op(e2, e4))) | (spl140_246 | ~ spl140_301)), inference(forward_demodulation, [], [f2594, f2892])).
fof(f2594, plain, (~ (e2 = op(e4, op(e4, e2))) | spl140_246), inference(avatar_component_clause, [], [f2593])).
fof(f5245, plain, (~ spl140_43 | spl140_83 | ~ spl140_315), inference(avatar_contradiction_clause, [], [f5244])).
fof(f5244, plain, ($false | (~ spl140_43 | spl140_83 | ~ spl140_315)), inference(subsumption_resolution, [], [f5243, f1583])).
fof(f1583, plain, (~ (e2 = op(e1, e3)) | spl140_83), inference(avatar_component_clause, [], [f1582])).
fof(f5243, plain, ((e2 = op(e1, e3)) | (~ spl140_43 | ~ spl140_315)), inference(forward_demodulation, [], [f2953, f1416])).
fof(f5228, plain, (~ spl140_1 | ~ spl140_16), inference(avatar_split_clause, [], [f5224, f1301, f1238])).
fof(f5224, plain, (~ (op(e4, e4) = unit) | ~ spl140_16), inference(backward_demodulation, [], [f474, f1303])).
fof(f1303, plain, ((op(e4, e1) = unit) | ~ spl140_16), inference(avatar_component_clause, [], [f1301])).
fof(f5222, plain, (~ spl140_1 | ~ spl140_26), inference(avatar_split_clause, [], [f5219, f1343, f1238])).
fof(f5219, plain, (~ (op(e4, e4) = unit) | ~ spl140_26), inference(backward_demodulation, [], [f427, f1345])).
fof(f5218, plain, (~ spl140_27 | ~ spl140_32), inference(avatar_split_clause, [], [f5215, f1368, f1347])).
fof(f5215, plain, (~ (e1 = op(e3, e4)) | ~ spl140_32), inference(backward_demodulation, [], [f467, f1370])).
fof(f1370, plain, ((e1 = op(e3, e3)) | ~ spl140_32), inference(avatar_component_clause, [], [f1368])).
fof(f5184, plain, (~ spl140_36 | ~ spl140_61), inference(avatar_split_clause, [], [f5181, f1490, f1385])).
fof(f5181, plain, (~ (op(e3, e2) = unit) | ~ spl140_61), inference(backward_demodulation, [], [f405, f1492])).
fof(f5177, plain, (~ spl140_19 | ~ spl140_69), inference(avatar_split_clause, [], [f5169, f1523, f1313])).
fof(f5169, plain, (~ (e3 = op(e4, e1)) | ~ spl140_69), inference(backward_demodulation, [], [f396, f1525])).
fof(f5150, plain, (spl140_52 | ~ spl140_12 | ~ spl140_301), inference(avatar_split_clause, [], [f5149, f2890, f1284, f1452])).
fof(f5149, plain, ((e1 = op(e2, e4)) | (~ spl140_12 | ~ spl140_301)), inference(forward_demodulation, [], [f2892, f1286])).
fof(f1286, plain, ((e1 = op(e4, e2)) | ~ spl140_12), inference(avatar_component_clause, [], [f1284])).
fof(f5106, plain, (~ spl140_8 | ~ spl140_54 | spl140_351), inference(avatar_contradiction_clause, [], [f5105])).
fof(f5105, plain, ($false | (~ spl140_8 | ~ spl140_54 | spl140_351)), inference(subsumption_resolution, [], [f5101, f1269])).
fof(f5101, plain, (~ (e2 = op(e4, e3)) | (~ spl140_54 | spl140_351)), inference(backward_demodulation, [], [f3193, f1462])).
fof(f3193, plain, (~ (e2 = op(e4, op(e2, e4))) | spl140_351), inference(avatar_component_clause, [], [f3191])).
fof(f5102, plain, (~ spl140_4 | ~ spl140_54), inference(avatar_split_clause, [], [f5097, f1460, f1250])).
fof(f5097, plain, (~ (e3 = op(e4, e4)) | ~ spl140_54), inference(backward_demodulation, [], [f426, f1462])).
fof(f426, plain, ~ (op(e2, e4) = op(e4, e4)), inference(cnf_transformation, [], [f5])).
fof(f5091, plain, (~ spl140_66 | ~ spl140_97 | ~ spl140_260), inference(avatar_contradiction_clause, [], [f5090])).
fof(f5090, plain, ($false | (~ spl140_66 | ~ spl140_97 | ~ spl140_260)), inference(subsumption_resolution, [], [f5089, f482])).
fof(f482, plain, ~ (e1 = e2), inference(cnf_transformation, [], [f6])).
fof(f5089, plain, ((e1 = e2) | (~ spl140_66 | ~ spl140_97 | ~ spl140_260)), inference(forward_demodulation, [], [f5083, f1643])).
fof(f5083, plain, ((e2 = op(e1, unit)) | (~ spl140_66 | ~ spl140_260)), inference(backward_demodulation, [], [f2656, f1513])).
fof(f2656, plain, ((e2 = op(e1, op(e2, e1))) | ~ spl140_260), inference(avatar_component_clause, [], [f2655])).
fof(f5086, plain, (~ spl140_51 | ~ spl140_66), inference(avatar_split_clause, [], [f5081, f1511, f1448])).
fof(f5081, plain, (~ (op(e2, e4) = unit) | ~ spl140_66), inference(backward_demodulation, [], [f454, f1513])).
fof(f5085, plain, (~ spl140_16 | ~ spl140_66), inference(avatar_split_clause, [], [f5079, f1511, f1301])).
fof(f5079, plain, (~ (op(e4, e1) = unit) | ~ spl140_66), inference(backward_demodulation, [], [f396, f1513])).
fof(f5069, plain, (~ spl140_64 | ~ spl140_89), inference(avatar_split_clause, [], [f5061, f1607, f1502])).
fof(f5061, plain, (~ (e3 = op(e2, e2)) | ~ spl140_89), inference(backward_demodulation, [], [f402, f1609])).
fof(f5059, plain, (~ spl140_70 | ~ spl140_95), inference(avatar_split_clause, [], [f5054, f1632, f1527])).
fof(f5054, plain, (~ (e4 = op(e2, e1)) | ~ spl140_95), inference(backward_demodulation, [], [f392, f1634])).
fof(f5048, plain, (spl140_37 | ~ spl140_43 | ~ spl140_257), inference(avatar_contradiction_clause, [], [f5047])).
fof(f5047, plain, ($false | (spl140_37 | ~ spl140_43 | ~ spl140_257)), inference(subsumption_resolution, [], [f5046, f1390])).
fof(f1390, plain, (~ (e1 = op(e3, e2)) | spl140_37), inference(avatar_component_clause, [], [f1389])).
fof(f5046, plain, ((e1 = op(e3, e2)) | (~ spl140_43 | ~ spl140_257)), inference(forward_demodulation, [], [f2644, f1416])).
fof(f2644, plain, ((e1 = op(e3, op(e3, e1))) | ~ spl140_257), inference(avatar_component_clause, [], [f2642])).
fof(f2642, plain, (spl140_257 <=> (e1 = op(e3, op(e3, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl140_257])])).
fof(f5032, plain, (spl140_19 | ~ spl140_27 | ~ spl140_349), inference(avatar_split_clause, [], [f5031, f3182, f1347, f1313])).
fof(f5031, plain, ((e3 = op(e4, e1)) | (~ spl140_27 | ~ spl140_349)), inference(forward_demodulation, [], [f3183, f1349])).
fof(f5008, plain, (~ spl140_12 | spl140_18 | ~ spl140_246), inference(avatar_contradiction_clause, [], [f5007])).
fof(f5007, plain, ($false | (~ spl140_12 | spl140_18 | ~ spl140_246)), inference(subsumption_resolution, [], [f5006, f1310])).
fof(f1310, plain, (~ (e2 = op(e4, e1)) | spl140_18), inference(avatar_component_clause, [], [f1309])).
fof(f5006, plain, ((e2 = op(e4, e1)) | (~ spl140_12 | ~ spl140_246)), inference(forward_demodulation, [], [f2595, f1286])).
fof(f2595, plain, ((e2 = op(e4, op(e4, e2))) | ~ spl140_246), inference(avatar_component_clause, [], [f2593])).
fof(f5001, plain, (~ spl140_2 | ~ spl140_12), inference(avatar_split_clause, [], [f4993, f1284, f1242])).
fof(f4993, plain, (~ (e1 = op(e4, e4)) | ~ spl140_12), inference(backward_demodulation, [], [f476, f1286])).
fof(f4957, plain, (~ spl140_65 | ~ spl140_70), inference(avatar_split_clause, [], [f4951, f1527, f1506])).
fof(f4951, plain, (~ (e4 = op(e2, e2)) | ~ spl140_70), inference(backward_demodulation, [], [f452, f1529])).
fof(f452, plain, ~ (op(e2, e1) = op(e2, e2)), inference(cnf_transformation, [], [f5])).
fof(f4956, plain, (~ spl140_45 | ~ spl140_70), inference(avatar_split_clause, [], [f4950, f1527, f1422])).
fof(f4950, plain, (~ (e4 = op(e3, e1)) | ~ spl140_70), inference(backward_demodulation, [], [f395, f1529])).
fof(f4949, plain, (spl140_12 | ~ spl140_78 | ~ spl140_353), inference(avatar_split_clause, [], [f4945, f3200, f1561, f1284])).
fof(f4945, plain, ((e1 = op(e4, e2)) | (~ spl140_78 | ~ spl140_353)), inference(backward_demodulation, [], [f3201, f1563])).
fof(f4948, plain, (~ spl140_78 | spl140_90 | ~ spl140_352), inference(avatar_contradiction_clause, [], [f4947])).
fof(f4947, plain, ($false | (~ spl140_78 | spl140_90 | ~ spl140_352)), inference(subsumption_resolution, [], [f4944, f1612])).
fof(f1612, plain, (~ (e4 = op(e1, e2)) | spl140_90), inference(avatar_component_clause, [], [f1611])).
fof(f4944, plain, ((e4 = op(e1, e2)) | (~ spl140_78 | ~ spl140_352)), inference(backward_demodulation, [], [f3198, f1563])).
fof(f3198, plain, ((e4 = op(e1, op(e1, e4))) | ~ spl140_352), inference(avatar_component_clause, [], [f3196])).
fof(f4899, plain, (~ spl140_113 | ~ spl140_114), inference(avatar_contradiction_clause, [], [f4898])).
fof(f4898, plain, ($false | (~ spl140_113 | ~ spl140_114)), inference(subsumption_resolution, [], [f4897, f485])).
fof(f4897, plain, ((e2 = e3) | (~ spl140_113 | ~ spl140_114)), inference(forward_demodulation, [], [f1714, f1710])).
fof(f1714, plain, ((e3 = op(unit, e2)) | ~ spl140_114), inference(avatar_component_clause, [], [f1712])).
fof(f1712, plain, (spl140_114 <=> (e3 = op(unit, e2))), introduced(avatar_definition, [new_symbols(naming, [spl140_114])])).
fof(f4864, plain, (~ spl140_32 | ~ spl140_37), inference(avatar_split_clause, [], [f4856, f1389, f1368])).
fof(f4856, plain, (~ (e1 = op(e3, e3)) | ~ spl140_37), inference(backward_demodulation, [], [f465, f1391])).
fof(f4863, plain, (~ spl140_12 | ~ spl140_37), inference(avatar_split_clause, [], [f4855, f1389, f1284])).
fof(f4855, plain, (~ (e1 = op(e4, e2)) | ~ spl140_37), inference(backward_demodulation, [], [f407, f1391])).
fof(f4844, plain, (~ spl140_1 | ~ spl140_51), inference(avatar_split_clause, [], [f4840, f1448, f1238])).
fof(f4840, plain, (~ (op(e4, e4) = unit) | ~ spl140_51), inference(backward_demodulation, [], [f426, f1450])).
fof(f4837, plain, (~ spl140_52 | ~ spl140_57), inference(avatar_split_clause, [], [f4827, f1473, f1452])).
fof(f4827, plain, (~ (e1 = op(e2, e4)) | ~ spl140_57), inference(backward_demodulation, [], [f457, f1475])).
fof(f4818, plain, (~ spl140_64 | ~ spl140_69), inference(avatar_split_clause, [], [f4809, f1523, f1502])).
fof(f4809, plain, (~ (e3 = op(e2, e2)) | ~ spl140_69), inference(backward_demodulation, [], [f452, f1525])).
fof(f4806, plain, (~ spl140_73 | ~ spl140_74), inference(avatar_contradiction_clause, [], [f4805])).
fof(f4805, plain, ($false | (~ spl140_73 | ~ spl140_74)), inference(subsumption_resolution, [], [f4804, f485])).
fof(f4804, plain, ((e2 = e3) | (~ spl140_73 | ~ spl140_74)), inference(forward_demodulation, [], [f1546, f1542])).
fof(f1542, plain, ((e2 = op(e2, unit)) | ~ spl140_73), inference(avatar_component_clause, [], [f1540])).
fof(f1540, plain, (spl140_73 <=> (e2 = op(e2, unit))), introduced(avatar_definition, [new_symbols(naming, [spl140_73])])).
fof(f1546, plain, ((e3 = op(e2, unit)) | ~ spl140_74), inference(avatar_component_clause, [], [f1544])).
fof(f1544, plain, (spl140_74 <=> (e3 = op(e2, unit))), introduced(avatar_definition, [new_symbols(naming, [spl140_74])])).
fof(f4802, plain, (~ spl140_121 | ~ spl140_125), inference(avatar_contradiction_clause, [], [f4801])).
fof(f4801, plain, ($false | (~ spl140_121 | ~ spl140_125)), inference(subsumption_resolution, [], [f4800, f1030])).
fof(f1030, plain, ~ (e4 = unit), inference(definition_unfolding, [], [f481, f377])).
fof(f481, plain, ~ (e0 = e4), inference(cnf_transformation, [], [f6])).
fof(f4800, plain, ((e4 = unit) | (~ spl140_121 | ~ spl140_125)), inference(backward_demodulation, [], [f1744, f1760])).
fof(f1760, plain, ((e4 = op(unit, unit)) | ~ spl140_125), inference(avatar_component_clause, [], [f1758])).
fof(f1758, plain, (spl140_125 <=> (e4 = op(unit, unit))), introduced(avatar_definition, [new_symbols(naming, [spl140_125])])).
fof(f1744, plain, ((unit = op(unit, unit)) | ~ spl140_121), inference(avatar_component_clause, [], [f1742])).
fof(f1742, plain, (spl140_121 <=> (unit = op(unit, unit))), introduced(avatar_definition, [new_symbols(naming, [spl140_121])])).
fof(f4771, plain, (~ spl140_45 | ~ spl140_85 | spl140_315), inference(avatar_split_clause, [], [f4551, f2951, f1590, f1422])).
fof(f4551, plain, (~ (e4 = op(e3, e1)) | (~ spl140_85 | spl140_315)), inference(forward_demodulation, [], [f2952, f1592])).
fof(f2952, plain, (~ (op(e1, e3) = op(e3, e1)) | spl140_315), inference(avatar_component_clause, [], [f2951])).
fof(f4769, plain, (~ spl140_36 | ~ spl140_86), inference(avatar_split_clause, [], [f4768, f1595, f1385])).
fof(f4768, plain, (~ (op(e3, e2) = unit) | ~ spl140_86), inference(forward_demodulation, [], [f403, f1597])).
fof(f4766, plain, (~ spl140_40 | ~ spl140_28 | spl140_348), inference(avatar_split_clause, [], [f4746, f3178, f1351, f1401])).
fof(f4746, plain, (~ (e4 = op(e3, e2)) | (~ spl140_28 | spl140_348)), inference(backward_demodulation, [], [f3179, f1353])).
fof(f3179, plain, (~ (e4 = op(e3, op(e3, e4))) | spl140_348), inference(avatar_component_clause, [], [f3178])).
fof(f4731, plain, (~ spl140_31 | ~ spl140_41), inference(avatar_split_clause, [], [f4725, f1406, f1364])).
fof(f4725, plain, (~ (op(e3, e3) = unit) | ~ spl140_41), inference(backward_demodulation, [], [f463, f1408])).
fof(f4730, plain, (~ spl140_16 | ~ spl140_41), inference(avatar_split_clause, [], [f4724, f1406, f1301])).
fof(f4724, plain, (~ (op(e4, e1) = unit) | ~ spl140_41), inference(backward_demodulation, [], [f397, f1408])).
fof(f4706, plain, (~ spl140_31 | ~ spl140_56), inference(avatar_split_clause, [], [f4700, f1469, f1364])).
fof(f4700, plain, (~ (op(e3, e3) = unit) | ~ spl140_56), inference(backward_demodulation, [], [f415, f1471])).
fof(f4699, plain, (~ spl140_73 | ~ spl140_75), inference(avatar_contradiction_clause, [], [f4698])).
fof(f4698, plain, ($false | (~ spl140_73 | ~ spl140_75)), inference(subsumption_resolution, [], [f4697, f486])).
fof(f486, plain, ~ (e2 = e4), inference(cnf_transformation, [], [f6])).
fof(f4697, plain, ((e2 = e4) | (~ spl140_73 | ~ spl140_75)), inference(forward_demodulation, [], [f1550, f1542])).
fof(f1550, plain, ((e4 = op(e2, unit)) | ~ spl140_75), inference(avatar_component_clause, [], [f1548])).
fof(f1548, plain, (spl140_75 <=> (e4 = op(e2, unit))), introduced(avatar_definition, [new_symbols(naming, [spl140_75])])).
fof(f4670, plain, (~ spl140_54 | ~ spl140_40 | spl140_249), inference(avatar_split_clause, [], [f4669, f2606, f1401, f1460])).
fof(f4669, plain, (~ (e3 = op(e2, e4)) | (~ spl140_40 | spl140_249)), inference(forward_demodulation, [], [f2608, f1403])).
fof(f4639, plain, (~ spl140_52 | ~ spl140_12 | spl140_301), inference(avatar_split_clause, [], [f4309, f2890, f1284, f1452])).
fof(f4309, plain, (~ (e1 = op(e2, e4)) | (~ spl140_12 | spl140_301)), inference(forward_demodulation, [], [f2891, f1286])).
fof(f2891, plain, (~ (op(e2, e4) = op(e4, e2)) | spl140_301), inference(avatar_component_clause, [], [f2890])).
fof(f4638, plain, (~ spl140_54 | ~ spl140_64), inference(avatar_split_clause, [], [f4637, f1502, f1460])).
fof(f4637, plain, (~ (e3 = op(e2, e4)) | ~ spl140_64), inference(forward_demodulation, [], [f456, f1504])).
fof(f4631, plain, (~ spl140_28 | ~ spl140_8 | spl140_298), inference(avatar_split_clause, [], [f4630, f2877, f1267, f1351])).
fof(f4630, plain, (~ (e2 = op(e3, e4)) | (~ spl140_8 | spl140_298)), inference(forward_demodulation, [], [f2878, f1269])).
fof(f2878, plain, (~ (op(e3, e4) = op(e4, e3)) | spl140_298), inference(avatar_component_clause, [], [f2877])).
fof(f4605, plain, (~ spl140_33 | ~ spl140_43), inference(avatar_split_clause, [], [f4598, f1414, f1372])).
fof(f4598, plain, (~ (e2 = op(e3, e3)) | ~ spl140_43), inference(backward_demodulation, [], [f463, f1416])).
fof(f4596, plain, (~ spl140_51 | ~ spl140_54), inference(avatar_contradiction_clause, [], [f4595])).
fof(f4595, plain, ($false | (~ spl140_51 | ~ spl140_54)), inference(subsumption_resolution, [], [f4593, f1031])).
fof(f1031, plain, ~ (e3 = unit), inference(definition_unfolding, [], [f480, f377])).
fof(f480, plain, ~ (e0 = e3), inference(cnf_transformation, [], [f6])).
fof(f4593, plain, ((e3 = unit) | (~ spl140_51 | ~ spl140_54)), inference(backward_demodulation, [], [f1450, f1462])).
fof(f4582, plain, (spl140_52 | ~ spl140_70 | ~ spl140_259), inference(avatar_contradiction_clause, [], [f4581])).
fof(f4581, plain, ($false | (spl140_52 | ~ spl140_70 | ~ spl140_259)), inference(subsumption_resolution, [], [f4580, f1453])).
fof(f1453, plain, (~ (e1 = op(e2, e4)) | spl140_52), inference(avatar_component_clause, [], [f1452])).
fof(f4580, plain, ((e1 = op(e2, e4)) | (~ spl140_70 | ~ spl140_259)), inference(forward_demodulation, [], [f2653, f1529])).
fof(f2653, plain, ((e1 = op(e2, op(e2, e1))) | ~ spl140_259), inference(avatar_component_clause, [], [f2651])).
fof(f2651, plain, (spl140_259 <=> (e1 = op(e2, op(e2, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl140_259])])).
fof(f4558, plain, (~ spl140_90 | ~ spl140_85), inference(avatar_split_clause, [], [f4411, f1590, f1611])).
fof(f4411, plain, (~ (e4 = op(e1, e2)) | ~ spl140_85), inference(forward_demodulation, [], [f445, f1592])).
fof(f4557, plain, (~ spl140_91 | ~ spl140_86), inference(avatar_split_clause, [], [f4556, f1595, f1616])).
fof(f4556, plain, (~ (op(e1, e1) = unit) | ~ spl140_86), inference(forward_demodulation, [], [f442, f1597])).
fof(f4553, plain, (~ spl140_90 | ~ spl140_78 | spl140_352), inference(avatar_split_clause, [], [f4515, f3196, f1561, f1611])).
fof(f4515, plain, (~ (e4 = op(e1, e2)) | (~ spl140_78 | spl140_352)), inference(backward_demodulation, [], [f3197, f1563])).
fof(f4538, plain, (~ spl140_3 | ~ spl140_18), inference(avatar_split_clause, [], [f4531, f1309, f1246])).
fof(f1246, plain, (spl140_3 <=> (e2 = op(e4, e4))), introduced(avatar_definition, [new_symbols(naming, [spl140_3])])).
fof(f4531, plain, (~ (e2 = op(e4, e4)) | ~ spl140_18), inference(backward_demodulation, [], [f474, f1311])).
fof(f4523, plain, (~ spl140_69 | ~ spl140_70), inference(avatar_contradiction_clause, [], [f4522])).
fof(f4522, plain, ($false | (~ spl140_69 | ~ spl140_70)), inference(subsumption_resolution, [], [f4521, f487])).
fof(f4521, plain, ((e3 = e4) | (~ spl140_69 | ~ spl140_70)), inference(backward_demodulation, [], [f1529, f1525])).
fof(f4518, plain, (~ spl140_18 | ~ spl140_78 | spl140_304), inference(avatar_split_clause, [], [f4513, f2903, f1561, f1309])).
fof(f4513, plain, (~ (e2 = op(e4, e1)) | (~ spl140_78 | spl140_304)), inference(backward_demodulation, [], [f2904, f1563])).
fof(f4517, plain, (~ spl140_3 | ~ spl140_78), inference(avatar_split_clause, [], [f4510, f1561, f1246])).
fof(f4510, plain, (~ (e2 = op(e4, e4)) | ~ spl140_78), inference(backward_demodulation, [], [f424, f1563])).
fof(f4491, plain, (~ spl140_78 | ~ spl140_70 | spl140_260), inference(avatar_split_clause, [], [f4220, f2655, f1527, f1561])).
fof(f4220, plain, (~ (e2 = op(e1, e4)) | (~ spl140_70 | spl140_260)), inference(forward_demodulation, [], [f2657, f1529])).
fof(f4480, plain, (~ spl140_18 | ~ spl140_93), inference(avatar_split_clause, [], [f4479, f1624, f1309])).
fof(f4479, plain, (~ (e2 = op(e4, e1)) | ~ spl140_93), inference(forward_demodulation, [], [f394, f1626])).
fof(f4475, plain, (~ spl140_12 | ~ spl140_70 | spl140_247), inference(avatar_contradiction_clause, [], [f4474])).
fof(f4474, plain, ($false | (~ spl140_12 | ~ spl140_70 | spl140_247)), inference(subsumption_resolution, [], [f4473, f1529])).
fof(f4473, plain, (~ (e4 = op(e2, e1)) | (~ spl140_12 | spl140_247)), inference(forward_demodulation, [], [f2599, f1286])).
fof(f2599, plain, (~ (e4 = op(e2, op(e4, e2))) | spl140_247), inference(avatar_component_clause, [], [f2597])).
fof(f4468, plain, (spl140_83 | ~ spl140_89 | ~ spl140_250), inference(avatar_contradiction_clause, [], [f4467])).
fof(f4467, plain, ($false | (spl140_83 | ~ spl140_89 | ~ spl140_250)), inference(subsumption_resolution, [], [f4466, f1583])).
fof(f4466, plain, ((e2 = op(e1, e3)) | (~ spl140_89 | ~ spl140_250)), inference(forward_demodulation, [], [f2613, f1609])).
fof(f4463, plain, (spl140_7 | ~ spl140_19 | ~ spl140_255), inference(avatar_contradiction_clause, [], [f4462])).
fof(f4462, plain, ($false | (spl140_7 | ~ spl140_19 | ~ spl140_255)), inference(subsumption_resolution, [], [f4461, f1264])).
fof(f1264, plain, (~ (e1 = op(e4, e3)) | spl140_7), inference(avatar_component_clause, [], [f1263])).
fof(f4461, plain, ((e1 = op(e4, e3)) | (~ spl140_19 | ~ spl140_255)), inference(forward_demodulation, [], [f2635, f1315])).
fof(f2635, plain, ((e1 = op(e4, op(e4, e1))) | ~ spl140_255), inference(avatar_component_clause, [], [f2633])).
fof(f4458, plain, (~ spl140_27 | spl140_45 | ~ spl140_348), inference(avatar_contradiction_clause, [], [f4457])).
fof(f4457, plain, ($false | (~ spl140_27 | spl140_45 | ~ spl140_348)), inference(subsumption_resolution, [], [f4456, f1423])).
fof(f1423, plain, (~ (e4 = op(e3, e1)) | spl140_45), inference(avatar_component_clause, [], [f1422])).
fof(f4456, plain, ((e4 = op(e3, e1)) | (~ spl140_27 | ~ spl140_348)), inference(forward_demodulation, [], [f3180, f1349])).
fof(f4455, plain, (spl140_79 | ~ spl140_85 | ~ spl140_241), inference(avatar_contradiction_clause, [], [f4454])).
fof(f4454, plain, ($false | (spl140_79 | ~ spl140_85 | ~ spl140_241)), inference(subsumption_resolution, [], [f4453, f1566])).
fof(f1566, plain, (~ (e3 = op(e1, e4)) | spl140_79), inference(avatar_component_clause, [], [f1565])).
fof(f4453, plain, ((e3 = op(e1, e4)) | (~ spl140_85 | ~ spl140_241)), inference(forward_demodulation, [], [f2573, f1592])).
fof(f4416, plain, (~ spl140_27 | ~ spl140_85 | spl140_242), inference(avatar_contradiction_clause, [], [f4415])).
fof(f4415, plain, ($false | (~ spl140_27 | ~ spl140_85 | spl140_242)), inference(subsumption_resolution, [], [f4414, f1349])).
fof(f4414, plain, (~ (e1 = op(e3, e4)) | (~ spl140_85 | spl140_242)), inference(forward_demodulation, [], [f2577, f1592])).
fof(f4406, plain, (~ spl140_19 | ~ spl140_85 | spl140_256), inference(avatar_contradiction_clause, [], [f4405])).
fof(f4405, plain, ($false | (~ spl140_19 | ~ spl140_85 | spl140_256)), inference(subsumption_resolution, [], [f4221, f1592])).
fof(f4221, plain, (~ (e4 = op(e1, e3)) | (~ spl140_19 | spl140_256)), inference(forward_demodulation, [], [f2639, f1315])).
fof(f2639, plain, (~ (e4 = op(e1, op(e4, e1))) | spl140_256), inference(avatar_component_clause, [], [f2637])).
fof(f4390, plain, (~ spl140_45 | ~ spl140_19 | spl140_305), inference(avatar_split_clause, [], [f4156, f2907, f1313, f1422])).
fof(f2907, plain, (spl140_305 <=> (e4 = op(op(e4, e1), e1))), introduced(avatar_definition, [new_symbols(naming, [spl140_305])])).
fof(f4156, plain, (~ (e4 = op(e3, e1)) | (~ spl140_19 | spl140_305)), inference(backward_demodulation, [], [f2909, f1315])).
fof(f2909, plain, (~ (e4 = op(op(e4, e1), e1)) | spl140_305), inference(avatar_component_clause, [], [f2907])).
fof(f4387, plain, (~ spl140_45 | ~ spl140_27 | spl140_348), inference(avatar_split_clause, [], [f4240, f3178, f1347, f1422])).
fof(f4240, plain, (~ (e4 = op(e3, e1)) | (~ spl140_27 | spl140_348)), inference(forward_demodulation, [], [f3179, f1349])).
fof(f4386, plain, (spl140_81 | ~ spl140_41 | ~ spl140_315), inference(avatar_split_clause, [], [f4385, f2951, f1406, f1574])).
fof(f4385, plain, ((op(e1, e3) = unit) | (~ spl140_41 | ~ spl140_315)), inference(forward_demodulation, [], [f2953, f1408])).
fof(f4382, plain, (~ spl140_91 | ~ spl140_41), inference(avatar_split_clause, [], [f4381, f1406, f1616])).
fof(f4381, plain, (~ (op(e1, e1) = unit) | ~ spl140_41), inference(forward_demodulation, [], [f393, f1408])).
fof(f4371, plain, (~ spl140_35 | ~ spl140_40), inference(avatar_split_clause, [], [f4368, f1401, f1380])).
fof(f4368, plain, (~ (e4 = op(e3, e3)) | ~ spl140_40), inference(backward_demodulation, [], [f465, f1403])).
fof(f4345, plain, (~ spl140_83 | ~ spl140_85), inference(avatar_contradiction_clause, [], [f4344])).
fof(f4344, plain, ($false | (~ spl140_83 | ~ spl140_85)), inference(subsumption_resolution, [], [f4343, f486])).
fof(f4343, plain, ((e2 = e4) | (~ spl140_83 | ~ spl140_85)), inference(forward_demodulation, [], [f1592, f1584])).
fof(f4329, plain, (~ spl140_97 | ~ spl140_99), inference(avatar_contradiction_clause, [], [f4328])).
fof(f4328, plain, ($false | (~ spl140_97 | ~ spl140_99)), inference(subsumption_resolution, [], [f4327, f483])).
fof(f4327, plain, ((e1 = e3) | (~ spl140_97 | ~ spl140_99)), inference(forward_demodulation, [], [f1651, f1643])).
fof(f1651, plain, ((e3 = op(e1, unit)) | ~ spl140_99), inference(avatar_component_clause, [], [f1649])).
fof(f1649, plain, (spl140_99 <=> (e3 = op(e1, unit))), introduced(avatar_definition, [new_symbols(naming, [spl140_99])])).
fof(f4302, plain, (~ spl140_90 | ~ spl140_70 | spl140_325), inference(avatar_split_clause, [], [f4301, f2995, f1527, f1611])).
fof(f4301, plain, (~ (e4 = op(e1, e2)) | (~ spl140_70 | spl140_325)), inference(forward_demodulation, [], [f2996, f1529])).
fof(f4284, plain, (spl140_79 | ~ spl140_19 | ~ spl140_304), inference(avatar_split_clause, [], [f4216, f2903, f1313, f1565])).
fof(f4216, plain, ((e3 = op(e1, e4)) | (~ spl140_19 | ~ spl140_304)), inference(forward_demodulation, [], [f2905, f1315])).
fof(f4262, plain, (~ spl140_89 | ~ spl140_43 | spl140_258), inference(avatar_split_clause, [], [f4261, f2646, f1414, f1607])).
fof(f2646, plain, (spl140_258 <=> (e3 = op(e1, op(e3, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl140_258])])).
fof(f4261, plain, (~ (e3 = op(e1, e2)) | (~ spl140_43 | spl140_258)), inference(forward_demodulation, [], [f2648, f1416])).
fof(f2648, plain, (~ (e3 = op(e1, op(e3, e1))) | spl140_258), inference(avatar_component_clause, [], [f2646])).
fof(f4237, plain, (spl140_37 | ~ spl140_89 | ~ spl140_335), inference(avatar_contradiction_clause, [], [f4236])).
fof(f4236, plain, ($false | (spl140_37 | ~ spl140_89 | ~ spl140_335)), inference(subsumption_resolution, [], [f4233, f1390])).
fof(f4233, plain, ((e1 = op(e3, e2)) | (~ spl140_89 | ~ spl140_335)), inference(backward_demodulation, [], [f3040, f1609])).
fof(f3040, plain, ((e1 = op(op(e1, e2), e2)) | ~ spl140_335), inference(avatar_component_clause, [], [f3039])).
fof(f4229, plain, (~ spl140_97 | ~ spl140_100), inference(avatar_contradiction_clause, [], [f4228])).
fof(f4228, plain, ($false | (~ spl140_97 | ~ spl140_100)), inference(subsumption_resolution, [], [f4227, f484])).
fof(f484, plain, ~ (e1 = e4), inference(cnf_transformation, [], [f6])).
fof(f4227, plain, ((e1 = e4) | (~ spl140_97 | ~ spl140_100)), inference(forward_demodulation, [], [f1655, f1643])).
fof(f1655, plain, ((e4 = op(e1, unit)) | ~ spl140_100), inference(avatar_component_clause, [], [f1653])).
fof(f1653, plain, (spl140_100 <=> (e4 = op(e1, unit))), introduced(avatar_definition, [new_symbols(naming, [spl140_100])])).
fof(f4225, plain, (~ spl140_113 | ~ spl140_115), inference(avatar_contradiction_clause, [], [f4224])).
fof(f4224, plain, ($false | (~ spl140_113 | ~ spl140_115)), inference(subsumption_resolution, [], [f4223, f486])).
fof(f4223, plain, ((e2 = e4) | (~ spl140_113 | ~ spl140_115)), inference(forward_demodulation, [], [f1718, f1710])).
fof(f1718, plain, ((e4 = op(unit, e2)) | ~ spl140_115), inference(avatar_component_clause, [], [f1716])).
fof(f1716, plain, (spl140_115 <=> (e4 = op(unit, e2))), introduced(avatar_definition, [new_symbols(naming, [spl140_115])])).
fof(f4188, plain, (spl140_90 | ~ spl140_70 | ~ spl140_325), inference(avatar_split_clause, [], [f4187, f2995, f1527, f1611])).
fof(f4187, plain, ((e4 = op(e1, e2)) | (~ spl140_70 | ~ spl140_325)), inference(forward_demodulation, [], [f2997, f1529])).
fof(f2997, plain, ((op(e1, e2) = op(e2, e1)) | ~ spl140_325), inference(avatar_component_clause, [], [f2995])).
fof(f4172, plain, (~ spl140_12 | spl140_78 | ~ spl140_303), inference(avatar_contradiction_clause, [], [f4171])).
fof(f4171, plain, ($false | (~ spl140_12 | spl140_78 | ~ spl140_303)), inference(subsumption_resolution, [], [f4169, f1562])).
fof(f4169, plain, ((e2 = op(e1, e4)) | (~ spl140_12 | ~ spl140_303)), inference(backward_demodulation, [], [f2900, f1286])).
fof(f4157, plain, (~ spl140_14 | ~ spl140_19), inference(avatar_split_clause, [], [f4150, f1313, f1292])).
fof(f4150, plain, (~ (e3 = op(e4, e2)) | ~ spl140_19), inference(backward_demodulation, [], [f472, f1315])).
fof(f472, plain, ~ (op(e4, e1) = op(e4, e2)), inference(cnf_transformation, [], [f5])).
fof(f4124, plain, (~ spl140_49 | ~ spl140_50), inference(avatar_contradiction_clause, [], [f4123])).
fof(f4123, plain, ($false | (~ spl140_49 | ~ spl140_50)), inference(subsumption_resolution, [], [f4122, f487])).
fof(f4122, plain, ((e3 = e4) | (~ spl140_49 | ~ spl140_50)), inference(forward_demodulation, [], [f1445, f1441])).
fof(f1445, plain, ((e4 = op(e3, unit)) | ~ spl140_50), inference(avatar_component_clause, [], [f1443])).
fof(f1443, plain, (spl140_50 <=> (e4 = op(e3, unit))), introduced(avatar_definition, [new_symbols(naming, [spl140_50])])).
fof(f4093, plain, (~ spl140_109 | ~ spl140_110), inference(avatar_contradiction_clause, [], [f4092])).
fof(f4092, plain, ($false | (~ spl140_109 | ~ spl140_110)), inference(subsumption_resolution, [], [f4091, f487])).
fof(f4091, plain, ((e3 = e4) | (~ spl140_109 | ~ spl140_110)), inference(forward_demodulation, [], [f1697, f1693])).
fof(f1693, plain, ((e3 = op(unit, e3)) | ~ spl140_109), inference(avatar_component_clause, [], [f1691])).
fof(f1691, plain, (spl140_109 <=> (e3 = op(unit, e3))), introduced(avatar_definition, [new_symbols(naming, [spl140_109])])).
fof(f1697, plain, ((e4 = op(unit, e3)) | ~ spl140_110), inference(avatar_component_clause, [], [f1695])).
fof(f1695, plain, (spl140_110 <=> (e4 = op(unit, e3))), introduced(avatar_definition, [new_symbols(naming, [spl140_110])])).
fof(f4087, plain, (spl140_12 | ~ spl140_90 | ~ spl140_335), inference(avatar_split_clause, [], [f4086, f3039, f1611, f1284])).
fof(f4086, plain, ((e1 = op(e4, e2)) | (~ spl140_90 | ~ spl140_335)), inference(forward_demodulation, [], [f3040, f1613])).
fof(f4071, plain, (~ spl140_57 | ~ spl140_83 | spl140_333), inference(avatar_split_clause, [], [f3853, f3030, f1582, f1473])).
fof(f3030, plain, (spl140_333 <=> (e1 = op(op(e1, e3), e3))), introduced(avatar_definition, [new_symbols(naming, [spl140_333])])).
fof(f3853, plain, (~ (e1 = op(e2, e3)) | (~ spl140_83 | spl140_333)), inference(backward_demodulation, [], [f3032, f1584])).
fof(f3032, plain, (~ (e1 = op(op(e1, e3), e3)) | spl140_333), inference(avatar_component_clause, [], [f3030])).
fof(f4066, plain, (~ spl140_26 | ~ spl140_36), inference(avatar_split_clause, [], [f4065, f1385, f1343])).
fof(f4065, plain, (~ (op(e3, e4) = unit) | ~ spl140_36), inference(forward_demodulation, [], [f466, f1387])).
fof(f1387, plain, ((op(e3, e2) = unit) | ~ spl140_36), inference(avatar_component_clause, [], [f1385])).
fof(f4060, plain, (~ spl140_6 | ~ spl140_8), inference(avatar_contradiction_clause, [], [f4059])).
fof(f4059, plain, ($false | (~ spl140_6 | ~ spl140_8)), inference(subsumption_resolution, [], [f4056, f1032])).
fof(f4056, plain, ((e2 = unit) | (~ spl140_6 | ~ spl140_8)), inference(backward_demodulation, [], [f1261, f1269])).
fof(f1261, plain, ((op(e4, e3) = unit) | ~ spl140_6), inference(avatar_component_clause, [], [f1259])).
fof(f4055, plain, (~ spl140_117 | ~ spl140_119), inference(avatar_contradiction_clause, [], [f4054])).
fof(f4054, plain, ($false | (~ spl140_117 | ~ spl140_119)), inference(subsumption_resolution, [], [f4053, f483])).
fof(f4053, plain, ((e1 = e3) | (~ spl140_117 | ~ spl140_119)), inference(forward_demodulation, [], [f1735, f1727])).
fof(f1735, plain, ((e3 = op(unit, e1)) | ~ spl140_119), inference(avatar_component_clause, [], [f1733])).
fof(f1733, plain, (spl140_119 <=> (e3 = op(unit, e1))), introduced(avatar_definition, [new_symbols(naming, [spl140_119])])).
fof(f4052, plain, (spl140_8 | ~ spl140_14 | ~ spl140_246), inference(avatar_split_clause, [], [f4051, f2593, f1292, f1267])).
fof(f4051, plain, ((e2 = op(e4, e3)) | (~ spl140_14 | ~ spl140_246)), inference(forward_demodulation, [], [f2595, f1294])).
fof(f4047, plain, (~ spl140_8 | ~ spl140_60 | spl140_323), inference(avatar_split_clause, [], [f3625, f2986, f1485, f1267])).
fof(f2986, plain, (spl140_323 <=> (e2 = op(op(e2, e3), e3))), introduced(avatar_definition, [new_symbols(naming, [spl140_323])])).
fof(f3625, plain, (~ (e2 = op(e4, e3)) | (~ spl140_60 | spl140_323)), inference(backward_demodulation, [], [f2988, f1487])).
fof(f2988, plain, (~ (e2 = op(op(e2, e3), e3)) | spl140_323), inference(avatar_component_clause, [], [f2986])).
fof(f4046, plain, (~ spl140_8 | ~ spl140_83), inference(avatar_split_clause, [], [f3850, f1582, f1267])).
fof(f3850, plain, (~ (e2 = op(e4, e3)) | ~ spl140_83), inference(backward_demodulation, [], [f414, f1584])).
fof(f414, plain, ~ (op(e1, e3) = op(e4, e3)), inference(cnf_transformation, [], [f5])).
fof(f4045, plain, (~ spl140_8 | ~ spl140_18), inference(avatar_split_clause, [], [f3891, f1309, f1267])).
fof(f3891, plain, (~ (e2 = op(e4, e3)) | ~ spl140_18), inference(backward_demodulation, [], [f473, f1311])).
fof(f473, plain, ~ (op(e4, e1) = op(e4, e3)), inference(cnf_transformation, [], [f5])).
fof(f4044, plain, (spl140_12 | ~ spl140_18 | ~ spl140_255), inference(avatar_contradiction_clause, [], [f4043])).
fof(f4043, plain, ($false | (spl140_12 | ~ spl140_18 | ~ spl140_255)), inference(subsumption_resolution, [], [f4042, f1285])).
fof(f1285, plain, (~ (e1 = op(e4, e2)) | spl140_12), inference(avatar_component_clause, [], [f1284])).
fof(f4042, plain, ((e1 = op(e4, e2)) | (~ spl140_18 | ~ spl140_255)), inference(forward_demodulation, [], [f2635, f1311])).
fof(f4040, plain, (spl140_27 | ~ spl140_45 | ~ spl140_257), inference(avatar_contradiction_clause, [], [f4039])).
fof(f4039, plain, ($false | (spl140_27 | ~ spl140_45 | ~ spl140_257)), inference(subsumption_resolution, [], [f4038, f1348])).
fof(f4038, plain, ((e1 = op(e3, e4)) | (~ spl140_45 | ~ spl140_257)), inference(forward_demodulation, [], [f2644, f1424])).
fof(f4036, plain, (spl140_57 | ~ spl140_69 | ~ spl140_259), inference(avatar_contradiction_clause, [], [f4035])).
fof(f4035, plain, ($false | (spl140_57 | ~ spl140_69 | ~ spl140_259)), inference(subsumption_resolution, [], [f4034, f1474])).
fof(f1474, plain, (~ (e1 = op(e2, e3)) | spl140_57), inference(avatar_component_clause, [], [f1473])).
fof(f4034, plain, ((e1 = op(e2, e3)) | (~ spl140_69 | ~ spl140_259)), inference(forward_demodulation, [], [f2653, f1525])).
fof(f4029, plain, (~ spl140_18 | ~ spl140_90 | spl140_256), inference(avatar_contradiction_clause, [], [f4028])).
fof(f4028, plain, ($false | (~ spl140_18 | ~ spl140_90 | spl140_256)), inference(subsumption_resolution, [], [f4027, f1613])).
fof(f4027, plain, (~ (e4 = op(e1, e2)) | (~ spl140_18 | spl140_256)), inference(forward_demodulation, [], [f2639, f1311])).
fof(f4026, plain, (~ spl140_45 | ~ spl140_79 | spl140_258), inference(avatar_contradiction_clause, [], [f4025])).
fof(f4025, plain, ($false | (~ spl140_45 | ~ spl140_79 | spl140_258)), inference(subsumption_resolution, [], [f4024, f1567])).
fof(f4024, plain, (~ (e3 = op(e1, e4)) | (~ spl140_45 | spl140_258)), inference(forward_demodulation, [], [f2648, f1424])).
fof(f4013, plain, (~ spl140_14 | ~ spl140_60 | spl140_247), inference(avatar_contradiction_clause, [], [f4012])).
fof(f4012, plain, ($false | (~ spl140_14 | ~ spl140_60 | spl140_247)), inference(subsumption_resolution, [], [f4011, f1487])).
fof(f4011, plain, (~ (e4 = op(e2, e3)) | (~ spl140_14 | spl140_247)), inference(forward_demodulation, [], [f2599, f1294])).
fof(f3999, plain, (~ spl140_2 | ~ spl140_25 | ~ spl140_264), inference(avatar_contradiction_clause, [], [f3998])).
fof(f3998, plain, ($false | (~ spl140_2 | ~ spl140_25 | ~ spl140_264)), inference(subsumption_resolution, [], [f3997, f1033])).
fof(f1033, plain, ~ (e1 = unit), inference(definition_unfolding, [], [f478, f377])).
fof(f478, plain, ~ (e0 = e1), inference(cnf_transformation, [], [f6])).
fof(f3997, plain, ((e1 = unit) | (~ spl140_2 | ~ spl140_25 | ~ spl140_264)), inference(forward_demodulation, [], [f3996, f1244])).
fof(f1244, plain, ((e1 = op(e4, e4)) | ~ spl140_2), inference(avatar_component_clause, [], [f1242])).
fof(f3996, plain, ((op(e4, e4) = unit) | (~ spl140_25 | ~ spl140_264)), inference(forward_demodulation, [], [f2675, f1340])).
fof(f3981, plain, (~ spl140_25 | ~ spl140_105 | spl140_265), inference(avatar_contradiction_clause, [], [f3980])).
fof(f3980, plain, ($false | (~ spl140_25 | ~ spl140_105 | spl140_265)), inference(subsumption_resolution, [], [f3979, f1676])).
fof(f1676, plain, ((e4 = op(unit, e4)) | ~ spl140_105), inference(avatar_component_clause, [], [f1674])).
fof(f1674, plain, (spl140_105 <=> (e4 = op(unit, e4))), introduced(avatar_definition, [new_symbols(naming, [spl140_105])])).
fof(f3979, plain, (~ (e4 = op(unit, e4)) | (~ spl140_25 | spl140_265)), inference(forward_demodulation, [], [f2679, f1340])).
fof(f2679, plain, (~ (e4 = op(unit, op(e4, unit))) | spl140_265), inference(avatar_component_clause, [], [f2677])).
fof(f2677, plain, (spl140_265 <=> (e4 = op(unit, op(e4, unit)))), introduced(avatar_definition, [new_symbols(naming, [spl140_265])])).
fof(f3978, plain, (~ spl140_49 | ~ spl140_109 | spl140_267), inference(avatar_contradiction_clause, [], [f3977])).
fof(f3977, plain, ($false | (~ spl140_49 | ~ spl140_109 | spl140_267)), inference(subsumption_resolution, [], [f3976, f1693])).
fof(f3976, plain, (~ (e3 = op(unit, e3)) | (~ spl140_49 | spl140_267)), inference(forward_demodulation, [], [f2688, f1441])).
fof(f2688, plain, (~ (e3 = op(unit, op(e3, unit))) | spl140_267), inference(avatar_component_clause, [], [f2686])).
fof(f2686, plain, (spl140_267 <=> (e3 = op(unit, op(e3, unit)))), introduced(avatar_definition, [new_symbols(naming, [spl140_267])])).
fof(f3972, plain, (~ spl140_97 | ~ spl140_117 | spl140_271), inference(avatar_contradiction_clause, [], [f3971])).
fof(f3971, plain, ($false | (~ spl140_97 | ~ spl140_117 | spl140_271)), inference(subsumption_resolution, [], [f3970, f1727])).
fof(f3970, plain, (~ (e1 = op(unit, e1)) | (~ spl140_97 | spl140_271)), inference(forward_demodulation, [], [f2706, f1643])).
fof(f2706, plain, (~ (e1 = op(unit, op(e1, unit))) | spl140_271), inference(avatar_component_clause, [], [f2704])).
fof(f2704, plain, (spl140_271 <=> (e1 = op(unit, op(e1, unit)))), introduced(avatar_definition, [new_symbols(naming, [spl140_271])])).
fof(f3969, plain, (spl140_54 | ~ spl140_60 | ~ spl140_239), inference(avatar_contradiction_clause, [], [f3968])).
fof(f3968, plain, ($false | (spl140_54 | ~ spl140_60 | ~ spl140_239)), inference(subsumption_resolution, [], [f3967, f1461])).
fof(f1461, plain, (~ (e3 = op(e2, e4)) | spl140_54), inference(avatar_component_clause, [], [f1460])).
fof(f3967, plain, ((e3 = op(e2, e4)) | (~ spl140_60 | ~ spl140_239)), inference(forward_demodulation, [], [f2564, f1487])).
fof(f3965, plain, (~ spl140_121 | ~ spl140_124), inference(avatar_contradiction_clause, [], [f3964])).
fof(f3964, plain, ($false | (~ spl140_121 | ~ spl140_124)), inference(subsumption_resolution, [], [f3962, f1031])).
fof(f3962, plain, ((e3 = unit) | (~ spl140_121 | ~ spl140_124)), inference(backward_demodulation, [], [f1744, f1756])).
fof(f1756, plain, ((e3 = op(unit, unit)) | ~ spl140_124), inference(avatar_component_clause, [], [f1754])).
fof(f1754, plain, (spl140_124 <=> (e3 = op(unit, unit))), introduced(avatar_definition, [new_symbols(naming, [spl140_124])])).
fof(f3942, plain, (~ spl140_121 | ~ spl140_123), inference(avatar_contradiction_clause, [], [f3941])).
fof(f3941, plain, ($false | (~ spl140_121 | ~ spl140_123)), inference(subsumption_resolution, [], [f3938, f1032])).
fof(f3938, plain, ((e2 = unit) | (~ spl140_121 | ~ spl140_123)), inference(backward_demodulation, [], [f1744, f1752])).
fof(f1752, plain, ((e2 = op(unit, unit)) | ~ spl140_123), inference(avatar_component_clause, [], [f1750])).
fof(f1750, plain, (spl140_123 <=> (e2 = op(unit, unit))), introduced(avatar_definition, [new_symbols(naming, [spl140_123])])).
fof(f3935, plain, (~ spl140_28 | spl140_40 | ~ spl140_348), inference(avatar_contradiction_clause, [], [f3934])).
fof(f3934, plain, ($false | (~ spl140_28 | spl140_40 | ~ spl140_348)), inference(subsumption_resolution, [], [f3933, f1402])).
fof(f1402, plain, (~ (e4 = op(e3, e2)) | spl140_40), inference(avatar_component_clause, [], [f1401])).
fof(f3933, plain, ((e4 = op(e3, e2)) | (~ spl140_28 | ~ spl140_348)), inference(forward_demodulation, [], [f3180, f1353])).
fof(f3917, plain, (~ spl140_11 | ~ spl140_6), inference(avatar_split_clause, [], [f3916, f1259, f1280])).
fof(f3916, plain, (~ (op(e4, e2) = unit) | ~ spl140_6), inference(forward_demodulation, [], [f475, f1261])).
fof(f3912, plain, (~ spl140_2 | ~ spl140_4), inference(avatar_contradiction_clause, [], [f3911])).
fof(f3911, plain, ($false | (~ spl140_2 | ~ spl140_4)), inference(subsumption_resolution, [], [f3910, f483])).
fof(f3910, plain, ((e1 = e3) | (~ spl140_2 | ~ spl140_4)), inference(backward_demodulation, [], [f1252, f1244])).
fof(f1252, plain, ((e3 = op(e4, e4)) | ~ spl140_4), inference(avatar_component_clause, [], [f1250])).
fof(f3887, plain, (~ spl140_3 | ~ spl140_28), inference(avatar_split_clause, [], [f3882, f1351, f1246])).
fof(f3882, plain, (~ (e2 = op(e4, e4)) | ~ spl140_28), inference(backward_demodulation, [], [f427, f1353])).
fof(f3881, plain, (~ spl140_31 | ~ spl140_32), inference(avatar_contradiction_clause, [], [f3880])).
fof(f3880, plain, ($false | (~ spl140_31 | ~ spl140_32)), inference(subsumption_resolution, [], [f3879, f1033])).
fof(f3879, plain, ((e1 = unit) | (~ spl140_31 | ~ spl140_32)), inference(forward_demodulation, [], [f1366, f1370])).
fof(f3878, plain, (~ spl140_7 | ~ spl140_32), inference(avatar_split_clause, [], [f3874, f1368, f1263])).
fof(f3874, plain, (~ (e1 = op(e4, e3)) | ~ spl140_32), inference(backward_demodulation, [], [f417, f1370])).
fof(f3872, plain, (~ spl140_51 | ~ spl140_73 | ~ spl140_350), inference(avatar_contradiction_clause, [], [f3871])).
fof(f3871, plain, ($false | (~ spl140_51 | ~ spl140_73 | ~ spl140_350)), inference(subsumption_resolution, [], [f3870, f486])).
fof(f3870, plain, ((e2 = e4) | (~ spl140_51 | ~ spl140_73 | ~ spl140_350)), inference(forward_demodulation, [], [f3868, f1542])).
fof(f3868, plain, ((e4 = op(e2, unit)) | (~ spl140_51 | ~ spl140_350)), inference(backward_demodulation, [], [f3189, f1450])).
fof(f3863, plain, (~ spl140_62 | ~ spl140_64), inference(avatar_contradiction_clause, [], [f3862])).
fof(f3862, plain, ($false | (~ spl140_62 | ~ spl140_64)), inference(subsumption_resolution, [], [f3861, f483])).
fof(f3861, plain, ((e1 = e3) | (~ spl140_62 | ~ spl140_64)), inference(backward_demodulation, [], [f1504, f1496])).
fof(f3837, plain, (~ spl140_7 | ~ spl140_79 | spl140_353), inference(avatar_split_clause, [], [f3836, f3200, f1565, f1263])).
fof(f3836, plain, (~ (e1 = op(e4, e3)) | (~ spl140_79 | spl140_353)), inference(forward_demodulation, [], [f3202, f1567])).
fof(f3816, plain, (~ spl140_79 | spl140_85 | ~ spl140_352), inference(avatar_contradiction_clause, [], [f3815])).
fof(f3815, plain, ($false | (~ spl140_79 | spl140_85 | ~ spl140_352)), inference(subsumption_resolution, [], [f3814, f1591])).
fof(f1591, plain, (~ (e4 = op(e1, e3)) | spl140_85), inference(avatar_component_clause, [], [f1590])).
fof(f3814, plain, ((e4 = op(e1, e3)) | (~ spl140_79 | ~ spl140_352)), inference(forward_demodulation, [], [f3198, f1567])).
fof(f3789, plain, (~ spl140_76 | ~ spl140_81), inference(avatar_split_clause, [], [f3786, f1574, f1553])).
fof(f3786, plain, (~ (op(e1, e4) = unit) | ~ spl140_81), inference(backward_demodulation, [], [f447, f1576])).
fof(f3788, plain, (~ spl140_6 | ~ spl140_81), inference(avatar_split_clause, [], [f3785, f1574, f1259])).
fof(f3785, plain, (~ (op(e4, e3) = unit) | ~ spl140_81), inference(backward_demodulation, [], [f414, f1576])).
fof(f3765, plain, (~ spl140_2 | ~ spl140_3), inference(avatar_contradiction_clause, [], [f3764])).
fof(f3764, plain, ($false | (~ spl140_2 | ~ spl140_3)), inference(subsumption_resolution, [], [f3763, f482])).
fof(f3763, plain, ((e1 = e2) | (~ spl140_2 | ~ spl140_3)), inference(backward_demodulation, [], [f1248, f1244])).
fof(f1248, plain, ((e2 = op(e4, e4)) | ~ spl140_3), inference(avatar_component_clause, [], [f1246])).
fof(f3732, plain, (~ spl140_57 | ~ spl140_60), inference(avatar_contradiction_clause, [], [f3731])).
fof(f3731, plain, ($false | (~ spl140_57 | ~ spl140_60)), inference(subsumption_resolution, [], [f3730, f484])).
fof(f3730, plain, ((e1 = e4) | (~ spl140_57 | ~ spl140_60)), inference(backward_demodulation, [], [f1487, f1475])).
fof(f3726, plain, (~ spl140_11 | ~ spl140_61), inference(avatar_split_clause, [], [f3723, f1490, f1280])).
fof(f3723, plain, (~ (op(e4, e2) = unit) | ~ spl140_61), inference(backward_demodulation, [], [f406, f1492])).
fof(f3708, plain, (~ spl140_12 | ~ spl140_90 | spl140_335), inference(avatar_split_clause, [], [f3706, f3039, f1611, f1284])).
fof(f3706, plain, (~ (e1 = op(e4, e2)) | (~ spl140_90 | spl140_335)), inference(backward_demodulation, [], [f3041, f1613])).
fof(f3700, plain, (~ spl140_97 | ~ spl140_98), inference(avatar_contradiction_clause, [], [f3699])).
fof(f3699, plain, ($false | (~ spl140_97 | ~ spl140_98)), inference(subsumption_resolution, [], [f3698, f482])).
fof(f3698, plain, ((e1 = e2) | (~ spl140_97 | ~ spl140_98)), inference(forward_demodulation, [], [f1647, f1643])).
fof(f1647, plain, ((e2 = op(e1, unit)) | ~ spl140_98), inference(avatar_component_clause, [], [f1645])).
fof(f1645, plain, (spl140_98 <=> (e2 = op(e1, unit))), introduced(avatar_definition, [new_symbols(naming, [spl140_98])])).
fof(f3697, plain, (~ spl140_117 | ~ spl140_118), inference(avatar_contradiction_clause, [], [f3696])).
fof(f3696, plain, ($false | (~ spl140_117 | ~ spl140_118)), inference(subsumption_resolution, [], [f3695, f482])).
fof(f3695, plain, ((e1 = e2) | (~ spl140_117 | ~ spl140_118)), inference(forward_demodulation, [], [f1731, f1727])).
fof(f1731, plain, ((e2 = op(unit, e1)) | ~ spl140_118), inference(avatar_component_clause, [], [f1729])).
fof(f1729, plain, (spl140_118 <=> (e2 = op(unit, e1))), introduced(avatar_definition, [new_symbols(naming, [spl140_118])])).
fof(f3684, plain, (~ spl140_7 | ~ spl140_85 | spl140_333), inference(avatar_split_clause, [], [f3683, f3030, f1590, f1263])).
fof(f3683, plain, (~ (e1 = op(e4, e3)) | (~ spl140_85 | spl140_333)), inference(forward_demodulation, [], [f3032, f1592])).
fof(f3666, plain, (~ spl140_3 | ~ spl140_4), inference(avatar_contradiction_clause, [], [f3665])).
fof(f3665, plain, ($false | (~ spl140_3 | ~ spl140_4)), inference(subsumption_resolution, [], [f3664, f485])).
fof(f3664, plain, ((e2 = e3) | (~ spl140_3 | ~ spl140_4)), inference(forward_demodulation, [], [f1252, f1248])).
fof(f3662, plain, (~ spl140_7 | ~ spl140_85 | spl140_299), inference(avatar_contradiction_clause, [], [f3661])).
fof(f3661, plain, ($false | (~ spl140_7 | ~ spl140_85 | spl140_299)), inference(subsumption_resolution, [], [f3659, f1592])).
fof(f3659, plain, (~ (e4 = op(e1, e3)) | (~ spl140_7 | spl140_299)), inference(backward_demodulation, [], [f2883, f1265])).
fof(f2883, plain, (~ (e4 = op(op(e4, e3), e3)) | spl140_299), inference(avatar_component_clause, [], [f2881])).
fof(f2881, plain, (spl140_299 <=> (e4 = op(op(e4, e3), e3))), introduced(avatar_definition, [new_symbols(naming, [spl140_299])])).
fof(f3656, plain, (~ spl140_11 | ~ spl140_12), inference(avatar_contradiction_clause, [], [f3655])).
fof(f3655, plain, ($false | (~ spl140_11 | ~ spl140_12)), inference(subsumption_resolution, [], [f3654, f1033])).
fof(f3654, plain, ((e1 = unit) | (~ spl140_11 | ~ spl140_12)), inference(forward_demodulation, [], [f1282, f1286])).
fof(f3633, plain, (~ spl140_51 | ~ spl140_52), inference(avatar_contradiction_clause, [], [f3632])).
fof(f3632, plain, ($false | (~ spl140_51 | ~ spl140_52)), inference(subsumption_resolution, [], [f3631, f1033])).
fof(f3631, plain, ((e1 = unit) | (~ spl140_51 | ~ spl140_52)), inference(forward_demodulation, [], [f1450, f1454])).
fof(f3622, plain, (~ spl140_57 | ~ spl140_62), inference(avatar_split_clause, [], [f3618, f1494, f1473])).
fof(f3618, plain, (~ (e1 = op(e2, e3)) | ~ spl140_62), inference(backward_demodulation, [], [f455, f1496])).
fof(f3603, plain, (~ spl140_81 | ~ spl140_85), inference(avatar_contradiction_clause, [], [f3602])).
fof(f3602, plain, ($false | (~ spl140_81 | ~ spl140_85)), inference(subsumption_resolution, [], [f3601, f1030])).
fof(f3601, plain, ((e4 = unit) | (~ spl140_81 | ~ spl140_85)), inference(backward_demodulation, [], [f1576, f1592])).
fof(f3596, plain, (~ spl140_89 | ~ spl140_94), inference(avatar_split_clause, [], [f3592, f1628, f1607])).
fof(f3592, plain, (~ (e3 = op(e1, e2)) | ~ spl140_94), inference(backward_demodulation, [], [f442, f1630])).
fof(f3595, plain, (~ spl140_19 | ~ spl140_94), inference(avatar_split_clause, [], [f3591, f1628, f1313])).
fof(f3591, plain, (~ (e3 = op(e4, e1)) | ~ spl140_94), inference(backward_demodulation, [], [f394, f1630])).
fof(f3583, plain, (~ spl140_117 | ~ spl140_120), inference(avatar_contradiction_clause, [], [f3582])).
fof(f3582, plain, ($false | (~ spl140_117 | ~ spl140_120)), inference(subsumption_resolution, [], [f3581, f484])).
fof(f3581, plain, ((e1 = e4) | (~ spl140_117 | ~ spl140_120)), inference(forward_demodulation, [], [f1739, f1727])).
fof(f1739, plain, ((e4 = op(unit, e1)) | ~ spl140_120), inference(avatar_component_clause, [], [f1737])).
fof(f1737, plain, (spl140_120 <=> (e4 = op(unit, e1))), introduced(avatar_definition, [new_symbols(naming, [spl140_120])])).
fof(f3530, plain, (~ spl140_18 | ~ spl140_52 | spl140_351), inference(avatar_split_clause, [], [f3529, f3191, f1452, f1309])).
fof(f3529, plain, (~ (e2 = op(e4, e1)) | (~ spl140_52 | spl140_351)), inference(forward_demodulation, [], [f3193, f1454])).
fof(f3518, plain, (~ spl140_7 | ~ spl140_12), inference(avatar_split_clause, [], [f3517, f1284, f1263])).
fof(f3517, plain, (~ (e1 = op(e4, e3)) | ~ spl140_12), inference(forward_demodulation, [], [f475, f1286])).
fof(f3513, plain, (~ spl140_90 | ~ spl140_12 | spl140_302), inference(avatar_split_clause, [], [f3512, f2894, f1284, f1611])).
fof(f2894, plain, (spl140_302 <=> (e4 = op(op(e4, e2), e2))), introduced(avatar_definition, [new_symbols(naming, [spl140_302])])).
fof(f3512, plain, (~ (e4 = op(e1, e2)) | (~ spl140_12 | spl140_302)), inference(forward_demodulation, [], [f2896, f1286])).
fof(f2896, plain, (~ (e4 = op(op(e4, e2), e2)) | spl140_302), inference(avatar_component_clause, [], [f2894])).
fof(f3500, plain, (~ spl140_12 | ~ spl140_14), inference(avatar_contradiction_clause, [], [f3499])).
fof(f3499, plain, ($false | (~ spl140_12 | ~ spl140_14)), inference(subsumption_resolution, [], [f3498, f483])).
fof(f3498, plain, ((e1 = e3) | (~ spl140_12 | ~ spl140_14)), inference(backward_demodulation, [], [f1294, f1286])).
fof(f3497, plain, (~ spl140_14 | ~ spl140_40 | spl140_302), inference(avatar_contradiction_clause, [], [f3496])).
fof(f3496, plain, ($false | (~ spl140_14 | ~ spl140_40 | spl140_302)), inference(subsumption_resolution, [], [f3493, f1403])).
fof(f3493, plain, (~ (e4 = op(e3, e2)) | (~ spl140_14 | spl140_302)), inference(backward_demodulation, [], [f2896, f1294])).
fof(f3475, plain, (~ spl140_20 | ~ spl140_25), inference(avatar_split_clause, [], [f3470, f1338, f1317])).
fof(f1317, plain, (spl140_20 <=> (e4 = op(e4, e1))), introduced(avatar_definition, [new_symbols(naming, [spl140_20])])).
fof(f3470, plain, (~ (e4 = op(e4, e1)) | ~ spl140_25), inference(backward_demodulation, [], [f981, f1340])).
fof(f981, plain, ~ (op(e4, e1) = op(e4, unit)), inference(definition_unfolding, [], [f468, f377])).
fof(f468, plain, ~ (op(e4, e0) = op(e4, e1)), inference(cnf_transformation, [], [f5])).
fof(f3474, plain, (~ spl140_15 | ~ spl140_25), inference(avatar_split_clause, [], [f3469, f1338, f1296])).
fof(f1296, plain, (spl140_15 <=> (e4 = op(e4, e2))), introduced(avatar_definition, [new_symbols(naming, [spl140_15])])).
fof(f3469, plain, (~ (e4 = op(e4, e2)) | ~ spl140_25), inference(backward_demodulation, [], [f980, f1340])).
fof(f980, plain, ~ (op(e4, e2) = op(e4, unit)), inference(definition_unfolding, [], [f469, f377])).
fof(f469, plain, ~ (op(e4, e0) = op(e4, e2)), inference(cnf_transformation, [], [f5])).
fof(f3473, plain, (~ spl140_10 | ~ spl140_25), inference(avatar_split_clause, [], [f3468, f1338, f1275])).
fof(f1275, plain, (spl140_10 <=> (e4 = op(e4, e3))), introduced(avatar_definition, [new_symbols(naming, [spl140_10])])).
fof(f3468, plain, (~ (e4 = op(e4, e3)) | ~ spl140_25), inference(backward_demodulation, [], [f979, f1340])).
fof(f979, plain, ~ (op(e4, e3) = op(e4, unit)), inference(definition_unfolding, [], [f470, f377])).
fof(f470, plain, ~ (op(e4, e0) = op(e4, e3)), inference(cnf_transformation, [], [f5])).
fof(f3460, plain, (~ spl140_32 | ~ spl140_33), inference(avatar_contradiction_clause, [], [f3459])).
fof(f3459, plain, ($false | (~ spl140_32 | ~ spl140_33)), inference(subsumption_resolution, [], [f3458, f482])).
fof(f3458, plain, ((e1 = e2) | (~ spl140_32 | ~ spl140_33)), inference(backward_demodulation, [], [f1374, f1370])).
fof(f3457, plain, (~ spl140_28 | ~ spl140_33), inference(avatar_split_clause, [], [f3453, f1372, f1351])).
fof(f3453, plain, (~ (e2 = op(e3, e4)) | ~ spl140_33), inference(backward_demodulation, [], [f467, f1374])).
fof(f3456, plain, (~ spl140_8 | ~ spl140_33), inference(avatar_split_clause, [], [f3452, f1372, f1267])).
fof(f3452, plain, (~ (e2 = op(e4, e3)) | ~ spl140_33), inference(backward_demodulation, [], [f417, f1374])).
fof(f3430, plain, (~ spl140_24 | ~ spl140_49), inference(avatar_split_clause, [], [f3424, f1439, f1334])).
fof(f1334, plain, (spl140_24 <=> (e3 = op(e4, unit))), introduced(avatar_definition, [new_symbols(naming, [spl140_24])])).
fof(f3424, plain, (~ (e3 = op(e4, unit)) | ~ spl140_49), inference(backward_demodulation, [], [f1020, f1441])).
fof(f1020, plain, ~ (op(e3, unit) = op(e4, unit)), inference(definition_unfolding, [], [f387, f377, f377])).
fof(f387, plain, ~ (op(e3, e0) = op(e4, e0)), inference(cnf_transformation, [], [f5])).
fof(f3429, plain, (~ spl140_44 | ~ spl140_49), inference(avatar_split_clause, [], [f3423, f1439, f1418])).
fof(f1418, plain, (spl140_44 <=> (e3 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl140_44])])).
fof(f3423, plain, (~ (e3 = op(e3, e1)) | ~ spl140_49), inference(backward_demodulation, [], [f985, f1441])).
fof(f985, plain, ~ (op(e3, e1) = op(e3, unit)), inference(definition_unfolding, [], [f458, f377])).
fof(f458, plain, ~ (op(e3, e0) = op(e3, e1)), inference(cnf_transformation, [], [f5])).
fof(f3428, plain, (~ spl140_39 | ~ spl140_49), inference(avatar_split_clause, [], [f3422, f1439, f1397])).
fof(f1397, plain, (spl140_39 <=> (e3 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl140_39])])).
fof(f3422, plain, (~ (e3 = op(e3, e2)) | ~ spl140_49), inference(backward_demodulation, [], [f984, f1441])).
fof(f984, plain, ~ (op(e3, e2) = op(e3, unit)), inference(definition_unfolding, [], [f459, f377])).
fof(f459, plain, ~ (op(e3, e0) = op(e3, e2)), inference(cnf_transformation, [], [f5])).
fof(f3426, plain, (~ spl140_29 | ~ spl140_49), inference(avatar_split_clause, [], [f3420, f1439, f1355])).
fof(f1355, plain, (spl140_29 <=> (e3 = op(e3, e4))), introduced(avatar_definition, [new_symbols(naming, [spl140_29])])).
fof(f3420, plain, (~ (e3 = op(e3, e4)) | ~ spl140_49), inference(backward_demodulation, [], [f982, f1441])).
fof(f982, plain, ~ (op(e3, e4) = op(e3, unit)), inference(definition_unfolding, [], [f461, f377])).
fof(f461, plain, ~ (op(e3, e0) = op(e3, e4)), inference(cnf_transformation, [], [f5])).
fof(f3419, plain, (~ spl140_52 | ~ spl140_54), inference(avatar_contradiction_clause, [], [f3418])).
fof(f3418, plain, ($false | (~ spl140_52 | ~ spl140_54)), inference(subsumption_resolution, [], [f3417, f483])).
fof(f3417, plain, ((e1 = e3) | (~ spl140_52 | ~ spl140_54)), inference(backward_demodulation, [], [f1462, f1454])).
fof(f3403, plain, (~ spl140_6 | ~ spl140_56), inference(avatar_split_clause, [], [f3399, f1469, f1259])).
fof(f3399, plain, (~ (op(e4, e3) = unit) | ~ spl140_56), inference(backward_demodulation, [], [f416, f1471])).
fof(f3386, plain, (~ spl140_61 | ~ spl140_66), inference(avatar_split_clause, [], [f3380, f1511, f1490])).
fof(f3380, plain, (~ (op(e2, e2) = unit) | ~ spl140_66), inference(backward_demodulation, [], [f452, f1513])).
fof(f3375, plain, (~ spl140_48 | ~ spl140_73), inference(avatar_split_clause, [], [f3368, f1540, f1435])).
fof(f1435, plain, (spl140_48 <=> (e2 = op(e3, unit))), introduced(avatar_definition, [new_symbols(naming, [spl140_48])])).
fof(f3368, plain, (~ (e2 = op(e3, unit)) | ~ spl140_73), inference(backward_demodulation, [], [f1022, f1542])).
fof(f1022, plain, ~ (op(e2, unit) = op(e3, unit)), inference(definition_unfolding, [], [f385, f377, f377])).
fof(f385, plain, ~ (op(e2, e0) = op(e3, e0)), inference(cnf_transformation, [], [f5])).
fof(f3374, plain, (~ spl140_23 | ~ spl140_73), inference(avatar_split_clause, [], [f3367, f1540, f1330])).
fof(f1330, plain, (spl140_23 <=> (e2 = op(e4, unit))), introduced(avatar_definition, [new_symbols(naming, [spl140_23])])).
fof(f3367, plain, (~ (e2 = op(e4, unit)) | ~ spl140_73), inference(backward_demodulation, [], [f1021, f1542])).
fof(f1021, plain, ~ (op(e2, unit) = op(e4, unit)), inference(definition_unfolding, [], [f386, f377, f377])).
fof(f386, plain, ~ (op(e2, e0) = op(e4, e0)), inference(cnf_transformation, [], [f5])).
fof(f3373, plain, (~ spl140_68 | ~ spl140_73), inference(avatar_split_clause, [], [f3366, f1540, f1519])).
fof(f1519, plain, (spl140_68 <=> (e2 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl140_68])])).
fof(f3366, plain, (~ (e2 = op(e2, e1)) | ~ spl140_73), inference(backward_demodulation, [], [f989, f1542])).
fof(f989, plain, ~ (op(e2, e1) = op(e2, unit)), inference(definition_unfolding, [], [f448, f377])).
fof(f448, plain, ~ (op(e2, e0) = op(e2, e1)), inference(cnf_transformation, [], [f5])).
fof(f3371, plain, (~ spl140_58 | ~ spl140_73), inference(avatar_split_clause, [], [f3364, f1540, f1477])).
fof(f1477, plain, (spl140_58 <=> (e2 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl140_58])])).
fof(f3364, plain, (~ (e2 = op(e2, e3)) | ~ spl140_73), inference(backward_demodulation, [], [f987, f1542])).
fof(f987, plain, ~ (op(e2, e3) = op(e2, unit)), inference(definition_unfolding, [], [f450, f377])).
fof(f450, plain, ~ (op(e2, e0) = op(e2, e3)), inference(cnf_transformation, [], [f5])).
fof(f3370, plain, (~ spl140_53 | ~ spl140_73), inference(avatar_split_clause, [], [f3363, f1540, f1456])).
fof(f1456, plain, (spl140_53 <=> (e2 = op(e2, e4))), introduced(avatar_definition, [new_symbols(naming, [spl140_53])])).
fof(f3363, plain, (~ (e2 = op(e2, e4)) | ~ spl140_73), inference(backward_demodulation, [], [f986, f1542])).
fof(f986, plain, ~ (op(e2, e4) = op(e2, unit)), inference(definition_unfolding, [], [f451, f377])).
fof(f451, plain, ~ (op(e2, e0) = op(e2, e4)), inference(cnf_transformation, [], [f5])).
fof(f3362, plain, (~ spl140_78 | ~ spl140_79), inference(avatar_contradiction_clause, [], [f3361])).
fof(f3361, plain, ($false | (~ spl140_78 | ~ spl140_79)), inference(subsumption_resolution, [], [f3360, f485])).
fof(f3360, plain, ((e2 = e3) | (~ spl140_78 | ~ spl140_79)), inference(backward_demodulation, [], [f1567, f1563])).
fof(f3331, plain, (~ spl140_61 | ~ spl140_86), inference(avatar_split_clause, [], [f3325, f1595, f1490])).
fof(f3325, plain, (~ (op(e2, e2) = unit) | ~ spl140_86), inference(backward_demodulation, [], [f402, f1597])).
fof(f3310, plain, (~ spl140_72 | ~ spl140_97), inference(avatar_split_clause, [], [f3302, f1641, f1536])).
fof(f1536, plain, (spl140_72 <=> (e1 = op(e2, unit))), introduced(avatar_definition, [new_symbols(naming, [spl140_72])])).
fof(f3302, plain, (~ (e1 = op(e2, unit)) | ~ spl140_97), inference(backward_demodulation, [], [f1025, f1643])).
fof(f1025, plain, ~ (op(e1, unit) = op(e2, unit)), inference(definition_unfolding, [], [f382, f377, f377])).
fof(f382, plain, ~ (op(e1, e0) = op(e2, e0)), inference(cnf_transformation, [], [f5])).
fof(f3309, plain, (~ spl140_47 | ~ spl140_97), inference(avatar_split_clause, [], [f3301, f1641, f1431])).
fof(f1431, plain, (spl140_47 <=> (e1 = op(e3, unit))), introduced(avatar_definition, [new_symbols(naming, [spl140_47])])).
fof(f3301, plain, (~ (e1 = op(e3, unit)) | ~ spl140_97), inference(backward_demodulation, [], [f1024, f1643])).
fof(f1024, plain, ~ (op(e1, unit) = op(e3, unit)), inference(definition_unfolding, [], [f383, f377, f377])).
fof(f383, plain, ~ (op(e1, e0) = op(e3, e0)), inference(cnf_transformation, [], [f5])).
fof(f3308, plain, (~ spl140_22 | ~ spl140_97), inference(avatar_split_clause, [], [f3300, f1641, f1326])).
fof(f3300, plain, (~ (e1 = op(e4, unit)) | ~ spl140_97), inference(backward_demodulation, [], [f1023, f1643])).
fof(f1023, plain, ~ (op(e1, unit) = op(e4, unit)), inference(definition_unfolding, [], [f384, f377, f377])).
fof(f384, plain, ~ (op(e1, e0) = op(e4, e0)), inference(cnf_transformation, [], [f5])).
fof(f3306, plain, (~ spl140_87 | ~ spl140_97), inference(avatar_split_clause, [], [f3298, f1641, f1599])).
fof(f1599, plain, (spl140_87 <=> (e1 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl140_87])])).
fof(f3298, plain, (~ (e1 = op(e1, e2)) | ~ spl140_97), inference(backward_demodulation, [], [f992, f1643])).
fof(f992, plain, ~ (op(e1, e2) = op(e1, unit)), inference(definition_unfolding, [], [f439, f377])).
fof(f439, plain, ~ (op(e1, e0) = op(e1, e2)), inference(cnf_transformation, [], [f5])).
fof(f3305, plain, (~ spl140_82 | ~ spl140_97), inference(avatar_split_clause, [], [f3297, f1641, f1578])).
fof(f1578, plain, (spl140_82 <=> (e1 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl140_82])])).
fof(f3297, plain, (~ (e1 = op(e1, e3)) | ~ spl140_97), inference(backward_demodulation, [], [f991, f1643])).
fof(f991, plain, ~ (op(e1, e3) = op(e1, unit)), inference(definition_unfolding, [], [f440, f377])).
fof(f440, plain, ~ (op(e1, e0) = op(e1, e3)), inference(cnf_transformation, [], [f5])).
fof(f3304, plain, (~ spl140_77 | ~ spl140_97), inference(avatar_split_clause, [], [f3296, f1641, f1557])).
fof(f1557, plain, (spl140_77 <=> (e1 = op(e1, e4))), introduced(avatar_definition, [new_symbols(naming, [spl140_77])])).
fof(f3296, plain, (~ (e1 = op(e1, e4)) | ~ spl140_97), inference(backward_demodulation, [], [f990, f1643])).
fof(f990, plain, ~ (op(e1, e4) = op(e1, unit)), inference(definition_unfolding, [], [f441, f377])).
fof(f441, plain, ~ (op(e1, e0) = op(e1, e4)), inference(cnf_transformation, [], [f5])).
fof(f3293, plain, (~ spl140_80 | ~ spl140_105), inference(avatar_split_clause, [], [f3287, f1674, f1569])).
fof(f1569, plain, (spl140_80 <=> (e4 = op(e1, e4))), introduced(avatar_definition, [new_symbols(naming, [spl140_80])])).
fof(f3287, plain, (~ (e4 = op(e1, e4)) | ~ spl140_105), inference(backward_demodulation, [], [f1007, f1676])).
fof(f1007, plain, ~ (op(e1, e4) = op(unit, e4)), inference(definition_unfolding, [], [f418, f377])).
fof(f418, plain, ~ (op(e0, e4) = op(e1, e4)), inference(cnf_transformation, [], [f5])).
fof(f3292, plain, (~ spl140_55 | ~ spl140_105), inference(avatar_split_clause, [], [f3286, f1674, f1464])).
fof(f1464, plain, (spl140_55 <=> (e4 = op(e2, e4))), introduced(avatar_definition, [new_symbols(naming, [spl140_55])])).
fof(f3286, plain, (~ (e4 = op(e2, e4)) | ~ spl140_105), inference(backward_demodulation, [], [f1006, f1676])).
fof(f1006, plain, ~ (op(e2, e4) = op(unit, e4)), inference(definition_unfolding, [], [f419, f377])).
fof(f419, plain, ~ (op(e0, e4) = op(e2, e4)), inference(cnf_transformation, [], [f5])).
fof(f3291, plain, (~ spl140_30 | ~ spl140_105), inference(avatar_split_clause, [], [f3285, f1674, f1359])).
fof(f1359, plain, (spl140_30 <=> (e4 = op(e3, e4))), introduced(avatar_definition, [new_symbols(naming, [spl140_30])])).
fof(f3285, plain, (~ (e4 = op(e3, e4)) | ~ spl140_105), inference(backward_demodulation, [], [f1005, f1676])).
fof(f1005, plain, ~ (op(e3, e4) = op(unit, e4)), inference(definition_unfolding, [], [f420, f377])).
fof(f420, plain, ~ (op(e0, e4) = op(e3, e4)), inference(cnf_transformation, [], [f5])).
fof(f3290, plain, (~ spl140_5 | ~ spl140_105), inference(avatar_split_clause, [], [f3284, f1674, f1254])).
fof(f1254, plain, (spl140_5 <=> (e4 = op(e4, e4))), introduced(avatar_definition, [new_symbols(naming, [spl140_5])])).
fof(f3284, plain, (~ (e4 = op(e4, e4)) | ~ spl140_105), inference(backward_demodulation, [], [f1004, f1676])).
fof(f1004, plain, ~ (op(e4, e4) = op(unit, e4)), inference(definition_unfolding, [], [f421, f377])).
fof(f421, plain, ~ (op(e0, e4) = op(e4, e4)), inference(cnf_transformation, [], [f5])).
fof(f3282, plain, (~ spl140_84 | ~ spl140_109), inference(avatar_split_clause, [], [f3276, f1691, f1586])).
fof(f1586, plain, (spl140_84 <=> (e3 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl140_84])])).
fof(f3276, plain, (~ (e3 = op(e1, e3)) | ~ spl140_109), inference(backward_demodulation, [], [f1011, f1693])).
fof(f1011, plain, ~ (op(e1, e3) = op(unit, e3)), inference(definition_unfolding, [], [f408, f377])).
fof(f408, plain, ~ (op(e0, e3) = op(e1, e3)), inference(cnf_transformation, [], [f5])).
fof(f3281, plain, (~ spl140_59 | ~ spl140_109), inference(avatar_split_clause, [], [f3275, f1691, f1481])).
fof(f1481, plain, (spl140_59 <=> (e3 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl140_59])])).
fof(f3275, plain, (~ (e3 = op(e2, e3)) | ~ spl140_109), inference(backward_demodulation, [], [f1010, f1693])).
fof(f1010, plain, ~ (op(e2, e3) = op(unit, e3)), inference(definition_unfolding, [], [f409, f377])).
fof(f409, plain, ~ (op(e0, e3) = op(e2, e3)), inference(cnf_transformation, [], [f5])).
fof(f3280, plain, (~ spl140_34 | ~ spl140_109), inference(avatar_split_clause, [], [f3274, f1691, f1376])).
fof(f1376, plain, (spl140_34 <=> (e3 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl140_34])])).
fof(f3274, plain, (~ (e3 = op(e3, e3)) | ~ spl140_109), inference(backward_demodulation, [], [f1009, f1693])).
fof(f1009, plain, ~ (op(e3, e3) = op(unit, e3)), inference(definition_unfolding, [], [f410, f377])).
fof(f410, plain, ~ (op(e0, e3) = op(e3, e3)), inference(cnf_transformation, [], [f5])).
fof(f3279, plain, (~ spl140_9 | ~ spl140_109), inference(avatar_split_clause, [], [f3273, f1691, f1271])).
fof(f1271, plain, (spl140_9 <=> (e3 = op(e4, e3))), introduced(avatar_definition, [new_symbols(naming, [spl140_9])])).
fof(f3273, plain, (~ (e3 = op(e4, e3)) | ~ spl140_109), inference(backward_demodulation, [], [f1008, f1693])).
fof(f1008, plain, ~ (op(e4, e3) = op(unit, e3)), inference(definition_unfolding, [], [f411, f377])).
fof(f411, plain, ~ (op(e0, e3) = op(e4, e3)), inference(cnf_transformation, [], [f5])).
fof(f3278, plain, (~ spl140_104 | ~ spl140_109), inference(avatar_split_clause, [], [f3272, f1691, f1670])).
fof(f1670, plain, (spl140_104 <=> (e3 = op(unit, e4))), introduced(avatar_definition, [new_symbols(naming, [spl140_104])])).
fof(f3272, plain, (~ (e3 = op(unit, e4)) | ~ spl140_109), inference(backward_demodulation, [], [f994, f1693])).
fof(f994, plain, ~ (op(unit, e3) = op(unit, e4)), inference(definition_unfolding, [], [f437, f377, f377])).
fof(f437, plain, ~ (op(e0, e3) = op(e0, e4)), inference(cnf_transformation, [], [f5])).
fof(f3270, plain, (~ spl140_88 | ~ spl140_113), inference(avatar_split_clause, [], [f3263, f1708, f1603])).
fof(f1603, plain, (spl140_88 <=> (e2 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl140_88])])).
fof(f3263, plain, (~ (e2 = op(e1, e2)) | ~ spl140_113), inference(backward_demodulation, [], [f1015, f1710])).
fof(f1015, plain, ~ (op(e1, e2) = op(unit, e2)), inference(definition_unfolding, [], [f398, f377])).
fof(f398, plain, ~ (op(e0, e2) = op(e1, e2)), inference(cnf_transformation, [], [f5])).
fof(f3269, plain, (~ spl140_63 | ~ spl140_113), inference(avatar_split_clause, [], [f3262, f1708, f1498])).
fof(f1498, plain, (spl140_63 <=> (e2 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl140_63])])).
fof(f3262, plain, (~ (e2 = op(e2, e2)) | ~ spl140_113), inference(backward_demodulation, [], [f1014, f1710])).
fof(f1014, plain, ~ (op(e2, e2) = op(unit, e2)), inference(definition_unfolding, [], [f399, f377])).
fof(f399, plain, ~ (op(e0, e2) = op(e2, e2)), inference(cnf_transformation, [], [f5])).
fof(f3268, plain, (~ spl140_38 | ~ spl140_113), inference(avatar_split_clause, [], [f3261, f1708, f1393])).
fof(f1393, plain, (spl140_38 <=> (e2 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl140_38])])).
fof(f3261, plain, (~ (e2 = op(e3, e2)) | ~ spl140_113), inference(backward_demodulation, [], [f1013, f1710])).
fof(f1013, plain, ~ (op(e3, e2) = op(unit, e2)), inference(definition_unfolding, [], [f400, f377])).
fof(f400, plain, ~ (op(e0, e2) = op(e3, e2)), inference(cnf_transformation, [], [f5])).
fof(f3267, plain, (~ spl140_13 | ~ spl140_113), inference(avatar_split_clause, [], [f3260, f1708, f1288])).
fof(f1288, plain, (spl140_13 <=> (e2 = op(e4, e2))), introduced(avatar_definition, [new_symbols(naming, [spl140_13])])).
fof(f3260, plain, (~ (e2 = op(e4, e2)) | ~ spl140_113), inference(backward_demodulation, [], [f1012, f1710])).
fof(f1012, plain, ~ (op(e4, e2) = op(unit, e2)), inference(definition_unfolding, [], [f401, f377])).
fof(f401, plain, ~ (op(e0, e2) = op(e4, e2)), inference(cnf_transformation, [], [f5])).
fof(f3266, plain, (~ spl140_108 | ~ spl140_113), inference(avatar_split_clause, [], [f3259, f1708, f1687])).
fof(f1687, plain, (spl140_108 <=> (e2 = op(unit, e3))), introduced(avatar_definition, [new_symbols(naming, [spl140_108])])).
fof(f3259, plain, (~ (e2 = op(unit, e3)) | ~ spl140_113), inference(backward_demodulation, [], [f996, f1710])).
fof(f996, plain, ~ (op(unit, e2) = op(unit, e3)), inference(definition_unfolding, [], [f435, f377, f377])).
fof(f435, plain, ~ (op(e0, e2) = op(e0, e3)), inference(cnf_transformation, [], [f5])).
fof(f3265, plain, (~ spl140_103 | ~ spl140_113), inference(avatar_split_clause, [], [f3258, f1708, f1666])).
fof(f1666, plain, (spl140_103 <=> (e2 = op(unit, e4))), introduced(avatar_definition, [new_symbols(naming, [spl140_103])])).
fof(f3258, plain, (~ (e2 = op(unit, e4)) | ~ spl140_113), inference(backward_demodulation, [], [f995, f1710])).
fof(f995, plain, ~ (op(unit, e2) = op(unit, e4)), inference(definition_unfolding, [], [f436, f377, f377])).
fof(f436, plain, ~ (op(e0, e2) = op(e0, e4)), inference(cnf_transformation, [], [f5])).
fof(f3256, plain, (~ spl140_92 | ~ spl140_117), inference(avatar_split_clause, [], [f3248, f1725, f1620])).
fof(f1620, plain, (spl140_92 <=> (e1 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl140_92])])).
fof(f3248, plain, (~ (e1 = op(e1, e1)) | ~ spl140_117), inference(backward_demodulation, [], [f1019, f1727])).
fof(f1019, plain, ~ (op(e1, e1) = op(unit, e1)), inference(definition_unfolding, [], [f388, f377])).
fof(f388, plain, ~ (op(e0, e1) = op(e1, e1)), inference(cnf_transformation, [], [f5])).
fof(f3255, plain, (~ spl140_67 | ~ spl140_117), inference(avatar_split_clause, [], [f3247, f1725, f1515])).
fof(f1515, plain, (spl140_67 <=> (e1 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl140_67])])).
fof(f3247, plain, (~ (e1 = op(e2, e1)) | ~ spl140_117), inference(backward_demodulation, [], [f1018, f1727])).
fof(f1018, plain, ~ (op(e2, e1) = op(unit, e1)), inference(definition_unfolding, [], [f389, f377])).
fof(f389, plain, ~ (op(e0, e1) = op(e2, e1)), inference(cnf_transformation, [], [f5])).
fof(f3254, plain, (~ spl140_42 | ~ spl140_117), inference(avatar_split_clause, [], [f3246, f1725, f1410])).
fof(f1410, plain, (spl140_42 <=> (e1 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl140_42])])).
fof(f3246, plain, (~ (e1 = op(e3, e1)) | ~ spl140_117), inference(backward_demodulation, [], [f1017, f1727])).
fof(f1017, plain, ~ (op(e3, e1) = op(unit, e1)), inference(definition_unfolding, [], [f390, f377])).
fof(f390, plain, ~ (op(e0, e1) = op(e3, e1)), inference(cnf_transformation, [], [f5])).
fof(f3253, plain, (~ spl140_17 | ~ spl140_117), inference(avatar_split_clause, [], [f3245, f1725, f1305])).
fof(f1305, plain, (spl140_17 <=> (e1 = op(e4, e1))), introduced(avatar_definition, [new_symbols(naming, [spl140_17])])).
fof(f3245, plain, (~ (e1 = op(e4, e1)) | ~ spl140_117), inference(backward_demodulation, [], [f1016, f1727])).
fof(f1016, plain, ~ (op(e4, e1) = op(unit, e1)), inference(definition_unfolding, [], [f391, f377])).
fof(f391, plain, ~ (op(e0, e1) = op(e4, e1)), inference(cnf_transformation, [], [f5])).
fof(f3252, plain, (~ spl140_112 | ~ spl140_117), inference(avatar_split_clause, [], [f3244, f1725, f1704])).
fof(f1704, plain, (spl140_112 <=> (e1 = op(unit, e2))), introduced(avatar_definition, [new_symbols(naming, [spl140_112])])).
fof(f3244, plain, (~ (e1 = op(unit, e2)) | ~ spl140_117), inference(backward_demodulation, [], [f999, f1727])).
fof(f999, plain, ~ (op(unit, e1) = op(unit, e2)), inference(definition_unfolding, [], [f432, f377, f377])).
fof(f432, plain, ~ (op(e0, e1) = op(e0, e2)), inference(cnf_transformation, [], [f5])).
fof(f3251, plain, (~ spl140_107 | ~ spl140_117), inference(avatar_split_clause, [], [f3243, f1725, f1683])).
fof(f1683, plain, (spl140_107 <=> (e1 = op(unit, e3))), introduced(avatar_definition, [new_symbols(naming, [spl140_107])])).
fof(f3243, plain, (~ (e1 = op(unit, e3)) | ~ spl140_117), inference(backward_demodulation, [], [f998, f1727])).
fof(f998, plain, ~ (op(unit, e1) = op(unit, e3)), inference(definition_unfolding, [], [f433, f377, f377])).
fof(f433, plain, ~ (op(e0, e1) = op(e0, e3)), inference(cnf_transformation, [], [f5])).
fof(f3250, plain, (~ spl140_102 | ~ spl140_117), inference(avatar_split_clause, [], [f3242, f1725, f1662])).
fof(f1662, plain, (spl140_102 <=> (e1 = op(unit, e4))), introduced(avatar_definition, [new_symbols(naming, [spl140_102])])).
fof(f3242, plain, (~ (e1 = op(unit, e4)) | ~ spl140_117), inference(backward_demodulation, [], [f997, f1727])).
fof(f997, plain, ~ (op(unit, e1) = op(unit, e4)), inference(definition_unfolding, [], [f434, f377, f377])).
fof(f434, plain, ~ (op(e0, e1) = op(e0, e4)), inference(cnf_transformation, [], [f5])).
fof(f3241, plain, (~ spl140_96 | ~ spl140_121), inference(avatar_split_clause, [], [f3233, f1742, f1637])).
fof(f1637, plain, (spl140_96 <=> (unit = op(e1, unit))), introduced(avatar_definition, [new_symbols(naming, [spl140_96])])).
fof(f3233, plain, (~ (unit = op(e1, unit)) | ~ spl140_121), inference(backward_demodulation, [], [f1029, f1744])).
fof(f1029, plain, ~ (op(e1, unit) = op(unit, unit)), inference(definition_unfolding, [], [f378, f377, f377, f377])).
fof(f378, plain, ~ (op(e0, e0) = op(e1, e0)), inference(cnf_transformation, [], [f5])).
fof(f3239, plain, (~ spl140_46 | ~ spl140_121), inference(avatar_split_clause, [], [f3231, f1742, f1427])).
fof(f1427, plain, (spl140_46 <=> (unit = op(e3, unit))), introduced(avatar_definition, [new_symbols(naming, [spl140_46])])).
fof(f3231, plain, (~ (unit = op(e3, unit)) | ~ spl140_121), inference(backward_demodulation, [], [f1027, f1744])).
fof(f1027, plain, ~ (op(e3, unit) = op(unit, unit)), inference(definition_unfolding, [], [f380, f377, f377, f377])).
fof(f380, plain, ~ (op(e0, e0) = op(e3, e0)), inference(cnf_transformation, [], [f5])).
fof(f3238, plain, (~ spl140_21 | ~ spl140_121), inference(avatar_split_clause, [], [f3230, f1742, f1322])).
fof(f1322, plain, (spl140_21 <=> (unit = op(e4, unit))), introduced(avatar_definition, [new_symbols(naming, [spl140_21])])).
fof(f3230, plain, (~ (unit = op(e4, unit)) | ~ spl140_121), inference(backward_demodulation, [], [f1026, f1744])).
fof(f1026, plain, ~ (op(e4, unit) = op(unit, unit)), inference(definition_unfolding, [], [f381, f377, f377, f377])).
fof(f381, plain, ~ (op(e0, e0) = op(e4, e0)), inference(cnf_transformation, [], [f5])).
fof(f3237, plain, (~ spl140_116 | ~ spl140_121), inference(avatar_split_clause, [], [f3229, f1742, f1721])).
fof(f1721, plain, (spl140_116 <=> (unit = op(unit, e1))), introduced(avatar_definition, [new_symbols(naming, [spl140_116])])).
fof(f3229, plain, (~ (unit = op(unit, e1)) | ~ spl140_121), inference(backward_demodulation, [], [f1003, f1744])).
fof(f1003, plain, ~ (op(unit, e1) = op(unit, unit)), inference(definition_unfolding, [], [f428, f377, f377, f377])).
fof(f428, plain, ~ (op(e0, e0) = op(e0, e1)), inference(cnf_transformation, [], [f5])).
fof(f3235, plain, (~ spl140_106 | ~ spl140_121), inference(avatar_split_clause, [], [f3227, f1742, f1679])).
fof(f1679, plain, (spl140_106 <=> (unit = op(unit, e3))), introduced(avatar_definition, [new_symbols(naming, [spl140_106])])).
fof(f3227, plain, (~ (unit = op(unit, e3)) | ~ spl140_121), inference(backward_demodulation, [], [f1001, f1744])).
fof(f1001, plain, ~ (op(unit, e3) = op(unit, unit)), inference(definition_unfolding, [], [f430, f377, f377, f377])).
fof(f430, plain, ~ (op(e0, e0) = op(e0, e3)), inference(cnf_transformation, [], [f5])).
fof(f3234, plain, (~ spl140_101 | ~ spl140_121), inference(avatar_split_clause, [], [f3226, f1742, f1658])).
fof(f1658, plain, (spl140_101 <=> (unit = op(unit, e4))), introduced(avatar_definition, [new_symbols(naming, [spl140_101])])).
fof(f3226, plain, (~ (unit = op(unit, e4)) | ~ spl140_121), inference(backward_demodulation, [], [f1000, f1744])).
fof(f1000, plain, ~ (op(unit, e4) = op(unit, unit)), inference(definition_unfolding, [], [f431, f377, f377, f377])).
fof(f431, plain, ~ (op(e0, e0) = op(e0, e4)), inference(cnf_transformation, [], [f5])).
fof(f3224, plain, (spl140_297 | spl140_296 | spl140_295 | spl140_294 | ~ spl140_95), inference(avatar_split_clause, [], [f793, f1632, f2841, f2850, f2859, f2868])).
fof(f2868, plain, (spl140_297 <=> sP0), introduced(avatar_definition, [new_symbols(naming, [spl140_297])])).
fof(f2859, plain, (spl140_296 <=> sP1), introduced(avatar_definition, [new_symbols(naming, [spl140_296])])).
fof(f2850, plain, (spl140_295 <=> sP2), introduced(avatar_definition, [new_symbols(naming, [spl140_295])])).
fof(f2841, plain, (spl140_294 <=> sP3), introduced(avatar_definition, [new_symbols(naming, [spl140_294])])).
fof(f793, plain, (~ (e4 = op(e1, e1)) | sP3 | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f150])).
fof(f150, plain, (((e4 = op(op(e4, e4), e4)) | ~ (e4 = op(op(e4, e4), e4)) | (op(e4, e4) = op(e4, e4))) & ((e3 = op(op(e4, e3), e4)) | ~ (e4 = op(op(e4, e3), e3)) | (op(e3, e4) = op(e4, e3))) & ((e2 = op(op(e4, e2), e4)) | ~ (e4 = op(op(e4, e2), e2)) | (op(e2, e4) = op(e4, e2))) & ((e1 = op(op(e4, e1), e4)) | ~ (e4 = op(op(e4, e1), e1)) | (op(e1, e4) = op(e4, e1))) & ((e0 = op(op(e4, e0), e4)) | ~ (e4 = op(op(e4, e0), e0)) | (op(e0, e4) = op(e4, e0))) & ((e4 = op(op(e3, e4), e3)) | ~ (e3 = op(op(e3, e4), e4)) | (op(e3, e4) = op(e4, e3))) & ((e3 = op(op(e3, e3), e3)) | ~ (e3 = op(op(e3, e3), e3)) | (op(e3, e3) = op(e3, e3))) & ((e2 = op(op(e3, e2), e3)) | ~ (e3 = op(op(e3, e2), e2)) | (op(e2, e3) = op(e3, e2))) & ((e1 = op(op(e3, e1), e3)) | ~ (e3 = op(op(e3, e1), e1)) | (op(e1, e3) = op(e3, e1))) & ((e0 = op(op(e3, e0), e3)) | ~ (e3 = op(op(e3, e0), e0)) | (op(e0, e3) = op(e3, e0))) & ((e4 = op(op(e2, e4), e2)) | ~ (e2 = op(op(e2, e4), e4)) | (op(e2, e4) = op(e4, e2))) & ((e3 = op(op(e2, e3), e2)) | ~ (e2 = op(op(e2, e3), e3)) | (op(e2, e3) = op(e3, e2))) & ((e2 = op(op(e2, e2), e2)) | ~ (e2 = op(op(e2, e2), e2)) | (op(e2, e2) = op(e2, e2))) & ((e1 = op(op(e2, e1), e2)) | ~ (e2 = op(op(e2, e1), e1)) | (op(e1, e2) = op(e2, e1))) & ((e0 = op(op(e2, e0), e2)) | ~ (e2 = op(op(e2, e0), e0)) | (op(e0, e2) = op(e2, e0))) & ((e4 = op(op(e1, e4), e1)) | ~ (e1 = op(op(e1, e4), e4)) | (op(e1, e4) = op(e4, e1))) & ((e3 = op(op(e1, e3), e1)) | ~ (e1 = op(op(e1, e3), e3)) | (op(e1, e3) = op(e3, e1))) & ((e2 = op(op(e1, e2), e1)) | ~ (e1 = op(op(e1, e2), e2)) | (op(e1, e2) = op(e2, e1))) & ((e1 = op(op(e1, e1), e1)) | ~ (e1 = op(op(e1, e1), e1)) | (op(e1, e1) = op(e1, e1))) & ((e0 = op(op(e1, e0), e1)) | ~ (e1 = op(op(e1, e0), e0)) | (op(e0, e1) = op(e1, e0))) & ((e4 = op(op(e0, e4), e0)) | ~ (e0 = op(op(e0, e4), e4)) | (op(e0, e4) = op(e4, e0))) & ((e3 = op(op(e0, e3), e0)) | ~ (e0 = op(op(e0, e3), e3)) | (op(e0, e3) = op(e3, e0))) & ((e2 = op(op(e0, e2), e0)) | ~ (e0 = op(op(e0, e2), e2)) | (op(e0, e2) = op(e2, e0))) & ((e1 = op(op(e0, e1), e0)) | ~ (e0 = op(op(e0, e1), e1)) | (op(e0, e1) = op(e1, e0))) & ((e0 = op(op(e0, e0), e0)) | ~ (e0 = op(op(e0, e0), e0)) | (op(e0, e0) = op(e0, e0))) & ((e4 = op(e4, e4)) | ~ (e4 = op(e4, e4)) | ~ (e4 = op(e4, e4))) & ((e4 = op(e4, e3)) | ~ (e4 = op(e3, e3)) | ~ (e3 = op(e4, e4))) & ((e4 = op(e4, e2)) | ~ (e4 = op(e2, e2)) | ~ (e2 = op(e4, e4))) & ((e4 = op(e4, e1)) | ~ (e4 = op(e1, e1)) | ~ (e1 = op(e4, e4))) & ((e4 = op(e4, e0)) | ~ (op(e0, e0) = e4) | ~ (e0 = op(e4, e4))) & ((e3 = op(e3, e4)) | ~ (e3 = op(e4, e4)) | ~ (e4 = op(e3, e3))) & ((e3 = op(e3, e3)) | ~ (e3 = op(e3, e3)) | ~ (e3 = op(e3, e3))) & ((e3 = op(e3, e2)) | ~ (e3 = op(e2, e2)) | ~ (e2 = op(e3, e3))) & ((e3 = op(e3, e1)) | ~ (e3 = op(e1, e1)) | ~ (e1 = op(e3, e3))) & ((e3 = op(e3, e0)) | ~ (op(e0, e0) = e3) | ~ (e0 = op(e3, e3))) & ((e2 = op(e2, e4)) | ~ (e2 = op(e4, e4)) | ~ (e4 = op(e2, e2))) & ((e2 = op(e2, e3)) | ~ (e2 = op(e3, e3)) | ~ (e3 = op(e2, e2))) & ((e2 = op(e2, e2)) | ~ (e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & ((e2 = op(e2, e1)) | ~ (e2 = op(e1, e1)) | ~ (e1 = op(e2, e2))) & ((e2 = op(e2, e0)) | ~ (op(e0, e0) = e2) | ~ (e0 = op(e2, e2))) & ((e1 = op(e1, e4)) | ~ (e1 = op(e4, e4)) | ~ (e4 = op(e1, e1))) & ((e1 = op(e1, e3)) | ~ (e1 = op(e3, e3)) | ~ (e3 = op(e1, e1))) & ((e1 = op(e1, e2)) | ~ (e1 = op(e2, e2)) | ~ (e2 = op(e1, e1))) & ((e1 = op(e1, e1)) | ~ (e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & ((e1 = op(e1, e0)) | ~ (op(e0, e0) = e1) | ~ (e0 = op(e1, e1))) & ((e0 = op(e0, e4)) | ~ (e0 = op(e4, e4)) | ~ (op(e0, e0) = e4)) & ((e0 = op(e0, e3)) | ~ (e0 = op(e3, e3)) | ~ (op(e0, e0) = e3)) & ((e0 = op(e0, e2)) | ~ (e0 = op(e2, e2)) | ~ (op(e0, e0) = e2)) & ((e0 = op(e0, e1)) | ~ (e0 = op(e1, e1)) | ~ (op(e0, e0) = e1)) & ((e0 = op(e0, e0)) | ~ (e0 = op(e0, e0)) | ~ (e0 = op(e0, e0))) & ((((e4 = unit) | ~ (e4 = op(e4, e4))) & (e4 = op(e4, e4))) | (((e4 = unit) | ~ (e3 = op(e4, e4))) & (e4 = op(e4, e3))) | (((e4 = unit) | ~ (e2 = op(e4, e4))) & (e4 = op(e4, e2))) | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118) & ((((e3 = unit) | ~ (e4 = op(e4, e3))) & (e3 = op(e4, e4))) | (((e3 = unit) | ~ (e3 = op(e4, e3))) & (e3 = op(e4, e3))) | (((e3 = unit) | ~ (e2 = op(e4, e3))) & (e3 = op(e4, e2))) | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96) & ((((e2 = unit) | ~ (e4 = op(e4, e2))) & (e2 = op(e4, e4))) | (((e2 = unit) | ~ (e3 = op(e4, e2))) & (e2 = op(e4, e3))) | (((e2 = unit) | ~ (e2 = op(e4, e2))) & (e2 = op(e4, e2))) | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74) & ((((e1 = unit) | ~ (e4 = op(e4, e1))) & (e1 = op(e4, e4))) | (((e1 = unit) | ~ (e3 = op(e4, e1))) & (e1 = op(e4, e3))) | (((e1 = unit) | ~ (e2 = op(e4, e1))) & (e1 = op(e4, e2))) | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52) & ((((e0 = unit) | ~ (e4 = op(e4, e0))) & (e0 = op(e4, e4))) | (((e0 = unit) | ~ (e3 = op(e4, e0))) & (e0 = op(e4, e3))) | (((e0 = unit) | ~ (e2 = op(e4, e0))) & (e0 = op(e4, e2))) | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30) & (((~ (e4 = op(e4, op(e4, e4))) | (e4 = op(e4, op(e4, e4)))) & (~ (e3 = op(e4, op(e3, e4))) | (e4 = op(e3, op(e3, e4)))) & (~ (e2 = op(e4, op(e2, e4))) | (e4 = op(e2, op(e2, e4)))) & (~ (e1 = op(e4, op(e1, e4))) | (e4 = op(e1, op(e1, e4)))) & (~ (e0 = op(e4, op(e0, e4))) | (e4 = op(e0, op(e0, e4))))) | sP29 | sP28 | sP27 | sP26) & ((~ (e4 = op(e4, e4)) & (e4 = op(e4, e4))) | (~ (e4 = op(e4, e3)) & (e3 = op(e4, e4))) | (~ (e4 = op(e4, e2)) & (e2 = op(e4, e4))) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4) & ((~ (e4 = op(e4, e4)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e1, e1)) & ~ (op(e0, e0) = e4)) | sP3 | sP2 | sP1 | sP0)), inference(definition_folding, [], [f9, e149, e148, e147, e146, e145, e144, e143, e142, e141, e140, e139, e138, e137, e136, e135, e134, e133, e132, e131, e130, e129, e128, e127, e126, e125, e124, e123, e122, e121, e120, e119, e118, e117, e116, e115, e114, e113, e112, e111, e110, e109, e108, e107, e106, e105, e104, e103, e102, e101, e100, e99, e98, e97, e96, e95, e94, e93, e92, e91, e90, e89, e88, e87, e86, e85, e84, e83, e82, e81, e80, e79, e78, e77, e76, e75, e74, e73, e72, e71, e70, e69, e68, e67, e66, e65, e64, e63, e62, e61, e60, e59, e58, e57, e56, e55, e54, e53, e52, e51, e50, e49, e48, e47, e46, e45, e44, e43, e42, e41, e40, e39, e38, e37, e36, e35, e34, e33, e32, e31, e30, e29, e28, e27, e26, e25, e24, e23, e22, e21, e20, e19, e18, e17, e16, e15, e14, e13, e12, e11, e10])).
fof(f10, plain, ((~ (e0 = op(e4, e4)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e0, e0))) | ~ sP0), inference(usedef, [], [e10])).
fof(e10, plain, (sP0 <=> (~ (e0 = op(e4, e4)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f11, plain, ((~ (e1 = op(e4, e4)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e1, e1)) & ~ (op(e0, e0) = e1)) | ~ sP1), inference(usedef, [], [e11])).
fof(e11, plain, (sP1 <=> (~ (e1 = op(e4, e4)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e1, e1)) & ~ (op(e0, e0) = e1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f12, plain, ((~ (e2 = op(e4, e4)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e1, e1)) & ~ (op(e0, e0) = e2)) | ~ sP2), inference(usedef, [], [e12])).
fof(e12, plain, (sP2 <=> (~ (e2 = op(e4, e4)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e1, e1)) & ~ (op(e0, e0) = e2))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f13, plain, ((~ (e3 = op(e4, e4)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e1, e1)) & ~ (op(e0, e0) = e3)) | ~ sP3), inference(usedef, [], [e13])).
fof(e13, plain, (sP3 <=> (~ (e3 = op(e4, e4)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e1, e1)) & ~ (op(e0, e0) = e3))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f14, plain, ((~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP4), inference(usedef, [], [e14])).
fof(e14, plain, (sP4 <=> (~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f15, plain, ((~ (e0 = op(e0, e1)) & (op(e0, e0) = e1)) | ~ sP5), inference(usedef, [], [e15])).
fof(e15, plain, (sP5 <=> (~ (e0 = op(e0, e1)) & (op(e0, e0) = e1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f16, plain, ((~ (e0 = op(e0, e2)) & (op(e0, e0) = e2)) | ~ sP6), inference(usedef, [], [e16])).
fof(e16, plain, (sP6 <=> (~ (e0 = op(e0, e2)) & (op(e0, e0) = e2))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f17, plain, ((~ (e0 = op(e0, e3)) & (op(e0, e0) = e3)) | ~ sP7), inference(usedef, [], [e17])).
fof(e17, plain, (sP7 <=> (~ (e0 = op(e0, e3)) & (op(e0, e0) = e3))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f18, plain, ((~ (e0 = op(e0, e4)) & (op(e0, e0) = e4)) | ~ sP8), inference(usedef, [], [e18])).
fof(e18, plain, (sP8 <=> (~ (e0 = op(e0, e4)) & (op(e0, e0) = e4))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f19, plain, ((~ (e1 = op(e1, e0)) & (e0 = op(e1, e1))) | ~ sP9), inference(usedef, [], [e19])).
fof(e19, plain, (sP9 <=> (~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f20, plain, ((~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP10), inference(usedef, [], [e20])).
fof(e20, plain, (sP10 <=> (~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f21, plain, ((~ (e1 = op(e1, e2)) & (e2 = op(e1, e1))) | ~ sP11), inference(usedef, [], [e21])).
fof(e21, plain, (sP11 <=> (~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f22, plain, ((~ (e1 = op(e1, e3)) & (e3 = op(e1, e1))) | ~ sP12), inference(usedef, [], [e22])).
fof(e22, plain, (sP12 <=> (~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f23, plain, ((~ (e1 = op(e1, e4)) & (e4 = op(e1, e1))) | ~ sP13), inference(usedef, [], [e23])).
fof(e23, plain, (sP13 <=> (~ (e1 = op(e1, e4)) & (e4 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f24, plain, ((~ (e2 = op(e2, e0)) & (e0 = op(e2, e2))) | ~ sP14), inference(usedef, [], [e24])).
fof(e24, plain, (sP14 <=> (~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f25, plain, ((~ (e2 = op(e2, e1)) & (e1 = op(e2, e2))) | ~ sP15), inference(usedef, [], [e25])).
fof(e25, plain, (sP15 <=> (~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f26, plain, ((~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP16), inference(usedef, [], [e26])).
fof(e26, plain, (sP16 <=> (~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f27, plain, ((~ (e2 = op(e2, e3)) & (e3 = op(e2, e2))) | ~ sP17), inference(usedef, [], [e27])).
fof(e27, plain, (sP17 <=> (~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f28, plain, ((~ (e2 = op(e2, e4)) & (e4 = op(e2, e2))) | ~ sP18), inference(usedef, [], [e28])).
fof(e28, plain, (sP18 <=> (~ (e2 = op(e2, e4)) & (e4 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f29, plain, ((~ (e3 = op(e3, e0)) & (e0 = op(e3, e3))) | ~ sP19), inference(usedef, [], [e29])).
fof(e29, plain, (sP19 <=> (~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP19])])).
fof(f30, plain, ((~ (e3 = op(e3, e1)) & (e1 = op(e3, e3))) | ~ sP20), inference(usedef, [], [e30])).
fof(e30, plain, (sP20 <=> (~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP20])])).
fof(f31, plain, ((~ (e3 = op(e3, e2)) & (e2 = op(e3, e3))) | ~ sP21), inference(usedef, [], [e31])).
fof(e31, plain, (sP21 <=> (~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP21])])).
fof(f32, plain, ((~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ~ sP22), inference(usedef, [], [e32])).
fof(e32, plain, (sP22 <=> (~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP22])])).
fof(f33, plain, ((~ (e3 = op(e3, e4)) & (e4 = op(e3, e3))) | ~ sP23), inference(usedef, [], [e33])).
fof(e33, plain, (sP23 <=> (~ (e3 = op(e3, e4)) & (e4 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP23])])).
fof(f34, plain, ((~ (e4 = op(e4, e0)) & (e0 = op(e4, e4))) | ~ sP24), inference(usedef, [], [e34])).
fof(e34, plain, (sP24 <=> (~ (e4 = op(e4, e0)) & (e0 = op(e4, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP24])])).
fof(f35, plain, ((~ (e4 = op(e4, e1)) & (e1 = op(e4, e4))) | ~ sP25), inference(usedef, [], [e35])).
fof(e35, plain, (sP25 <=> (~ (e4 = op(e4, e1)) & (e1 = op(e4, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP25])])).
fof(f36, plain, (((~ (e4 = op(e0, op(e4, e0))) | (e0 = op(e4, op(e4, e0)))) & (~ (e3 = op(e0, op(e3, e0))) | (e0 = op(e3, op(e3, e0)))) & (~ (e2 = op(e0, op(e2, e0))) | (e0 = op(e2, op(e2, e0)))) & (~ (e1 = op(e0, op(e1, e0))) | (e0 = op(e1, op(e1, e0)))) & (~ (e0 = op(e0, op(e0, e0))) | (e0 = op(e0, op(e0, e0))))) | ~ sP26), inference(usedef, [], [e36])).
fof(e36, plain, (sP26 <=> ((~ (e4 = op(e0, op(e4, e0))) | (e0 = op(e4, op(e4, e0)))) & (~ (e3 = op(e0, op(e3, e0))) | (e0 = op(e3, op(e3, e0)))) & (~ (e2 = op(e0, op(e2, e0))) | (e0 = op(e2, op(e2, e0)))) & (~ (e1 = op(e0, op(e1, e0))) | (e0 = op(e1, op(e1, e0)))) & (~ (e0 = op(e0, op(e0, e0))) | (e0 = op(e0, op(e0, e0)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP26])])).
fof(f37, plain, (((~ (e4 = op(e1, op(e4, e1))) | (e1 = op(e4, op(e4, e1)))) & (~ (e3 = op(e1, op(e3, e1))) | (e1 = op(e3, op(e3, e1)))) & (~ (e2 = op(e1, op(e2, e1))) | (e1 = op(e2, op(e2, e1)))) & (~ (e1 = op(e1, op(e1, e1))) | (e1 = op(e1, op(e1, e1)))) & (~ (e0 = op(e1, op(e0, e1))) | (e1 = op(e0, op(e0, e1))))) | ~ sP27), inference(usedef, [], [e37])).
fof(e37, plain, (sP27 <=> ((~ (e4 = op(e1, op(e4, e1))) | (e1 = op(e4, op(e4, e1)))) & (~ (e3 = op(e1, op(e3, e1))) | (e1 = op(e3, op(e3, e1)))) & (~ (e2 = op(e1, op(e2, e1))) | (e1 = op(e2, op(e2, e1)))) & (~ (e1 = op(e1, op(e1, e1))) | (e1 = op(e1, op(e1, e1)))) & (~ (e0 = op(e1, op(e0, e1))) | (e1 = op(e0, op(e0, e1)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP27])])).
fof(f38, plain, (((~ (e4 = op(e2, op(e4, e2))) | (e2 = op(e4, op(e4, e2)))) & (~ (e3 = op(e2, op(e3, e2))) | (e2 = op(e3, op(e3, e2)))) & (~ (e2 = op(e2, op(e2, e2))) | (e2 = op(e2, op(e2, e2)))) & (~ (e1 = op(e2, op(e1, e2))) | (e2 = op(e1, op(e1, e2)))) & (~ (e0 = op(e2, op(e0, e2))) | (e2 = op(e0, op(e0, e2))))) | ~ sP28), inference(usedef, [], [e38])).
fof(e38, plain, (sP28 <=> ((~ (e4 = op(e2, op(e4, e2))) | (e2 = op(e4, op(e4, e2)))) & (~ (e3 = op(e2, op(e3, e2))) | (e2 = op(e3, op(e3, e2)))) & (~ (e2 = op(e2, op(e2, e2))) | (e2 = op(e2, op(e2, e2)))) & (~ (e1 = op(e2, op(e1, e2))) | (e2 = op(e1, op(e1, e2)))) & (~ (e0 = op(e2, op(e0, e2))) | (e2 = op(e0, op(e0, e2)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP28])])).
fof(f39, plain, (((~ (e4 = op(e3, op(e4, e3))) | (e3 = op(e4, op(e4, e3)))) & (~ (e3 = op(e3, op(e3, e3))) | (e3 = op(e3, op(e3, e3)))) & (~ (e2 = op(e3, op(e2, e3))) | (e3 = op(e2, op(e2, e3)))) & (~ (e1 = op(e3, op(e1, e3))) | (e3 = op(e1, op(e1, e3)))) & (~ (e0 = op(e3, op(e0, e3))) | (e3 = op(e0, op(e0, e3))))) | ~ sP29), inference(usedef, [], [e39])).
fof(e39, plain, (sP29 <=> ((~ (e4 = op(e3, op(e4, e3))) | (e3 = op(e4, op(e4, e3)))) & (~ (e3 = op(e3, op(e3, e3))) | (e3 = op(e3, op(e3, e3)))) & (~ (e2 = op(e3, op(e2, e3))) | (e3 = op(e2, op(e2, e3)))) & (~ (e1 = op(e3, op(e1, e3))) | (e3 = op(e1, op(e1, e3)))) & (~ (e0 = op(e3, op(e0, e3))) | (e3 = op(e0, op(e0, e3)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP29])])).
fof(f40, plain, ((((e0 = unit) | ~ (e0 = op(e0, e0))) & (e0 = op(e0, e0))) | ~ sP30), inference(usedef, [], [e40])).
fof(e40, plain, (sP30 <=> (((e0 = unit) | ~ (e0 = op(e0, e0))) & (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP30])])).
fof(f41, plain, ((((e0 = unit) | ~ (op(e0, e0) = e1)) & (e0 = op(e0, e1))) | ~ sP31), inference(usedef, [], [e41])).
fof(e41, plain, (sP31 <=> (((e0 = unit) | ~ (op(e0, e0) = e1)) & (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP31])])).
fof(f42, plain, ((((e0 = unit) | ~ (op(e0, e0) = e2)) & (e0 = op(e0, e2))) | ~ sP32), inference(usedef, [], [e42])).
fof(e42, plain, (sP32 <=> (((e0 = unit) | ~ (op(e0, e0) = e2)) & (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP32])])).
fof(f43, plain, ((((e0 = unit) | ~ (op(e0, e0) = e3)) & (e0 = op(e0, e3))) | ~ sP33), inference(usedef, [], [e43])).
fof(e43, plain, (sP33 <=> (((e0 = unit) | ~ (op(e0, e0) = e3)) & (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP33])])).
fof(f44, plain, ((((e0 = unit) | ~ (op(e0, e0) = e4)) & (e0 = op(e0, e4))) | ~ sP34), inference(usedef, [], [e44])).
fof(e44, plain, (sP34 <=> (((e0 = unit) | ~ (op(e0, e0) = e4)) & (e0 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP34])])).
fof(f45, plain, ((((e0 = unit) | ~ (e0 = op(e1, e0))) & (e0 = op(e1, e0))) | ~ sP35), inference(usedef, [], [e45])).
fof(e45, plain, (sP35 <=> (((e0 = unit) | ~ (e0 = op(e1, e0))) & (e0 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP35])])).
fof(f46, plain, ((((e0 = unit) | ~ (e1 = op(e1, e0))) & (e0 = op(e1, e1))) | ~ sP36), inference(usedef, [], [e46])).
fof(e46, plain, (sP36 <=> (((e0 = unit) | ~ (e1 = op(e1, e0))) & (e0 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP36])])).
fof(f47, plain, ((((e0 = unit) | ~ (e2 = op(e1, e0))) & (e0 = op(e1, e2))) | ~ sP37), inference(usedef, [], [e47])).
fof(e47, plain, (sP37 <=> (((e0 = unit) | ~ (e2 = op(e1, e0))) & (e0 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP37])])).
fof(f48, plain, ((((e0 = unit) | ~ (e3 = op(e1, e0))) & (e0 = op(e1, e3))) | ~ sP38), inference(usedef, [], [e48])).
fof(e48, plain, (sP38 <=> (((e0 = unit) | ~ (e3 = op(e1, e0))) & (e0 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP38])])).
fof(f49, plain, ((((e0 = unit) | ~ (e4 = op(e1, e0))) & (e0 = op(e1, e4))) | ~ sP39), inference(usedef, [], [e49])).
fof(e49, plain, (sP39 <=> (((e0 = unit) | ~ (e4 = op(e1, e0))) & (e0 = op(e1, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP39])])).
fof(f50, plain, ((((e0 = unit) | ~ (e0 = op(e2, e0))) & (e0 = op(e2, e0))) | ~ sP40), inference(usedef, [], [e50])).
fof(e50, plain, (sP40 <=> (((e0 = unit) | ~ (e0 = op(e2, e0))) & (e0 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP40])])).
fof(f51, plain, ((((e0 = unit) | ~ (e1 = op(e2, e0))) & (e0 = op(e2, e1))) | ~ sP41), inference(usedef, [], [e51])).
fof(e51, plain, (sP41 <=> (((e0 = unit) | ~ (e1 = op(e2, e0))) & (e0 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP41])])).
fof(f52, plain, ((((e0 = unit) | ~ (e2 = op(e2, e0))) & (e0 = op(e2, e2))) | ~ sP42), inference(usedef, [], [e52])).
fof(e52, plain, (sP42 <=> (((e0 = unit) | ~ (e2 = op(e2, e0))) & (e0 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP42])])).
fof(f53, plain, ((((e0 = unit) | ~ (e3 = op(e2, e0))) & (e0 = op(e2, e3))) | ~ sP43), inference(usedef, [], [e53])).
fof(e53, plain, (sP43 <=> (((e0 = unit) | ~ (e3 = op(e2, e0))) & (e0 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP43])])).
fof(f54, plain, ((((e0 = unit) | ~ (e4 = op(e2, e0))) & (e0 = op(e2, e4))) | ~ sP44), inference(usedef, [], [e54])).
fof(e54, plain, (sP44 <=> (((e0 = unit) | ~ (e4 = op(e2, e0))) & (e0 = op(e2, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP44])])).
fof(f55, plain, ((((e0 = unit) | ~ (e0 = op(e3, e0))) & (e0 = op(e3, e0))) | ~ sP45), inference(usedef, [], [e55])).
fof(e55, plain, (sP45 <=> (((e0 = unit) | ~ (e0 = op(e3, e0))) & (e0 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP45])])).
fof(f56, plain, ((((e0 = unit) | ~ (e1 = op(e3, e0))) & (e0 = op(e3, e1))) | ~ sP46), inference(usedef, [], [e56])).
fof(e56, plain, (sP46 <=> (((e0 = unit) | ~ (e1 = op(e3, e0))) & (e0 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP46])])).
fof(f57, plain, ((((e0 = unit) | ~ (e2 = op(e3, e0))) & (e0 = op(e3, e2))) | ~ sP47), inference(usedef, [], [e57])).
fof(e57, plain, (sP47 <=> (((e0 = unit) | ~ (e2 = op(e3, e0))) & (e0 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP47])])).
fof(f58, plain, ((((e0 = unit) | ~ (e3 = op(e3, e0))) & (e0 = op(e3, e3))) | ~ sP48), inference(usedef, [], [e58])).
fof(e58, plain, (sP48 <=> (((e0 = unit) | ~ (e3 = op(e3, e0))) & (e0 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP48])])).
fof(f59, plain, ((((e0 = unit) | ~ (e4 = op(e3, e0))) & (e0 = op(e3, e4))) | ~ sP49), inference(usedef, [], [e59])).
fof(e59, plain, (sP49 <=> (((e0 = unit) | ~ (e4 = op(e3, e0))) & (e0 = op(e3, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP49])])).
fof(f60, plain, ((((e0 = unit) | ~ (e0 = op(e4, e0))) & (e0 = op(e4, e0))) | ~ sP50), inference(usedef, [], [e60])).
fof(e60, plain, (sP50 <=> (((e0 = unit) | ~ (e0 = op(e4, e0))) & (e0 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP50])])).
fof(f61, plain, ((((e0 = unit) | ~ (e1 = op(e4, e0))) & (e0 = op(e4, e1))) | ~ sP51), inference(usedef, [], [e61])).
fof(e61, plain, (sP51 <=> (((e0 = unit) | ~ (e1 = op(e4, e0))) & (e0 = op(e4, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP51])])).
fof(f62, plain, ((((e1 = unit) | ~ (e0 = op(e0, e1))) & (op(e0, e0) = e1)) | ~ sP52), inference(usedef, [], [e62])).
fof(e62, plain, (sP52 <=> (((e1 = unit) | ~ (e0 = op(e0, e1))) & (op(e0, e0) = e1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP52])])).
fof(f63, plain, ((((e1 = unit) | ~ (e1 = op(e0, e1))) & (e1 = op(e0, e1))) | ~ sP53), inference(usedef, [], [e63])).
fof(e63, plain, (sP53 <=> (((e1 = unit) | ~ (e1 = op(e0, e1))) & (e1 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP53])])).
fof(f64, plain, ((((e1 = unit) | ~ (e2 = op(e0, e1))) & (e1 = op(e0, e2))) | ~ sP54), inference(usedef, [], [e64])).
fof(e64, plain, (sP54 <=> (((e1 = unit) | ~ (e2 = op(e0, e1))) & (e1 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP54])])).
fof(f65, plain, ((((e1 = unit) | ~ (e3 = op(e0, e1))) & (e1 = op(e0, e3))) | ~ sP55), inference(usedef, [], [e65])).
fof(e65, plain, (sP55 <=> (((e1 = unit) | ~ (e3 = op(e0, e1))) & (e1 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP55])])).
fof(f66, plain, ((((e1 = unit) | ~ (e4 = op(e0, e1))) & (e1 = op(e0, e4))) | ~ sP56), inference(usedef, [], [e66])).
fof(e66, plain, (sP56 <=> (((e1 = unit) | ~ (e4 = op(e0, e1))) & (e1 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP56])])).
fof(f67, plain, ((((e1 = unit) | ~ (e0 = op(e1, e1))) & (e1 = op(e1, e0))) | ~ sP57), inference(usedef, [], [e67])).
fof(e67, plain, (sP57 <=> (((e1 = unit) | ~ (e0 = op(e1, e1))) & (e1 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP57])])).
fof(f68, plain, ((((e1 = unit) | ~ (e1 = op(e1, e1))) & (e1 = op(e1, e1))) | ~ sP58), inference(usedef, [], [e68])).
fof(e68, plain, (sP58 <=> (((e1 = unit) | ~ (e1 = op(e1, e1))) & (e1 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP58])])).
fof(f69, plain, ((((e1 = unit) | ~ (e2 = op(e1, e1))) & (e1 = op(e1, e2))) | ~ sP59), inference(usedef, [], [e69])).
fof(e69, plain, (sP59 <=> (((e1 = unit) | ~ (e2 = op(e1, e1))) & (e1 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP59])])).
fof(f70, plain, ((((e1 = unit) | ~ (e3 = op(e1, e1))) & (e1 = op(e1, e3))) | ~ sP60), inference(usedef, [], [e70])).
fof(e70, plain, (sP60 <=> (((e1 = unit) | ~ (e3 = op(e1, e1))) & (e1 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP60])])).
fof(f71, plain, ((((e1 = unit) | ~ (e4 = op(e1, e1))) & (e1 = op(e1, e4))) | ~ sP61), inference(usedef, [], [e71])).
fof(e71, plain, (sP61 <=> (((e1 = unit) | ~ (e4 = op(e1, e1))) & (e1 = op(e1, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP61])])).
fof(f72, plain, ((((e1 = unit) | ~ (e0 = op(e2, e1))) & (e1 = op(e2, e0))) | ~ sP62), inference(usedef, [], [e72])).
fof(e72, plain, (sP62 <=> (((e1 = unit) | ~ (e0 = op(e2, e1))) & (e1 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP62])])).
fof(f73, plain, ((((e1 = unit) | ~ (e1 = op(e2, e1))) & (e1 = op(e2, e1))) | ~ sP63), inference(usedef, [], [e73])).
fof(e73, plain, (sP63 <=> (((e1 = unit) | ~ (e1 = op(e2, e1))) & (e1 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP63])])).
fof(f74, plain, ((((e1 = unit) | ~ (e2 = op(e2, e1))) & (e1 = op(e2, e2))) | ~ sP64), inference(usedef, [], [e74])).
fof(e74, plain, (sP64 <=> (((e1 = unit) | ~ (e2 = op(e2, e1))) & (e1 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP64])])).
fof(f75, plain, ((((e1 = unit) | ~ (e3 = op(e2, e1))) & (e1 = op(e2, e3))) | ~ sP65), inference(usedef, [], [e75])).
fof(e75, plain, (sP65 <=> (((e1 = unit) | ~ (e3 = op(e2, e1))) & (e1 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP65])])).
fof(f76, plain, ((((e1 = unit) | ~ (e4 = op(e2, e1))) & (e1 = op(e2, e4))) | ~ sP66), inference(usedef, [], [e76])).
fof(e76, plain, (sP66 <=> (((e1 = unit) | ~ (e4 = op(e2, e1))) & (e1 = op(e2, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP66])])).
fof(f77, plain, ((((e1 = unit) | ~ (e0 = op(e3, e1))) & (e1 = op(e3, e0))) | ~ sP67), inference(usedef, [], [e77])).
fof(e77, plain, (sP67 <=> (((e1 = unit) | ~ (e0 = op(e3, e1))) & (e1 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP67])])).
fof(f78, plain, ((((e1 = unit) | ~ (e1 = op(e3, e1))) & (e1 = op(e3, e1))) | ~ sP68), inference(usedef, [], [e78])).
fof(e78, plain, (sP68 <=> (((e1 = unit) | ~ (e1 = op(e3, e1))) & (e1 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP68])])).
fof(f79, plain, ((((e1 = unit) | ~ (e2 = op(e3, e1))) & (e1 = op(e3, e2))) | ~ sP69), inference(usedef, [], [e79])).
fof(e79, plain, (sP69 <=> (((e1 = unit) | ~ (e2 = op(e3, e1))) & (e1 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP69])])).
fof(f80, plain, ((((e1 = unit) | ~ (e3 = op(e3, e1))) & (e1 = op(e3, e3))) | ~ sP70), inference(usedef, [], [e80])).
fof(e80, plain, (sP70 <=> (((e1 = unit) | ~ (e3 = op(e3, e1))) & (e1 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP70])])).
fof(f81, plain, ((((e1 = unit) | ~ (e4 = op(e3, e1))) & (e1 = op(e3, e4))) | ~ sP71), inference(usedef, [], [e81])).
fof(e81, plain, (sP71 <=> (((e1 = unit) | ~ (e4 = op(e3, e1))) & (e1 = op(e3, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP71])])).
fof(f82, plain, ((((e1 = unit) | ~ (e0 = op(e4, e1))) & (e1 = op(e4, e0))) | ~ sP72), inference(usedef, [], [e82])).
fof(e82, plain, (sP72 <=> (((e1 = unit) | ~ (e0 = op(e4, e1))) & (e1 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP72])])).
fof(f83, plain, ((((e1 = unit) | ~ (e1 = op(e4, e1))) & (e1 = op(e4, e1))) | ~ sP73), inference(usedef, [], [e83])).
fof(e83, plain, (sP73 <=> (((e1 = unit) | ~ (e1 = op(e4, e1))) & (e1 = op(e4, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP73])])).
fof(f84, plain, ((((e2 = unit) | ~ (e0 = op(e0, e2))) & (op(e0, e0) = e2)) | ~ sP74), inference(usedef, [], [e84])).
fof(e84, plain, (sP74 <=> (((e2 = unit) | ~ (e0 = op(e0, e2))) & (op(e0, e0) = e2))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP74])])).
fof(f85, plain, ((((e2 = unit) | ~ (e1 = op(e0, e2))) & (e2 = op(e0, e1))) | ~ sP75), inference(usedef, [], [e85])).
fof(e85, plain, (sP75 <=> (((e2 = unit) | ~ (e1 = op(e0, e2))) & (e2 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP75])])).
fof(f86, plain, ((((e2 = unit) | ~ (e2 = op(e0, e2))) & (e2 = op(e0, e2))) | ~ sP76), inference(usedef, [], [e86])).
fof(e86, plain, (sP76 <=> (((e2 = unit) | ~ (e2 = op(e0, e2))) & (e2 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP76])])).
fof(f87, plain, ((((e2 = unit) | ~ (e3 = op(e0, e2))) & (e2 = op(e0, e3))) | ~ sP77), inference(usedef, [], [e87])).
fof(e87, plain, (sP77 <=> (((e2 = unit) | ~ (e3 = op(e0, e2))) & (e2 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP77])])).
fof(f88, plain, ((((e2 = unit) | ~ (e4 = op(e0, e2))) & (e2 = op(e0, e4))) | ~ sP78), inference(usedef, [], [e88])).
fof(e88, plain, (sP78 <=> (((e2 = unit) | ~ (e4 = op(e0, e2))) & (e2 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP78])])).
fof(f89, plain, ((((e2 = unit) | ~ (e0 = op(e1, e2))) & (e2 = op(e1, e0))) | ~ sP79), inference(usedef, [], [e89])).
fof(e89, plain, (sP79 <=> (((e2 = unit) | ~ (e0 = op(e1, e2))) & (e2 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP79])])).
fof(f90, plain, ((((e2 = unit) | ~ (e1 = op(e1, e2))) & (e2 = op(e1, e1))) | ~ sP80), inference(usedef, [], [e90])).
fof(e90, plain, (sP80 <=> (((e2 = unit) | ~ (e1 = op(e1, e2))) & (e2 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP80])])).
fof(f91, plain, ((((e2 = unit) | ~ (e2 = op(e1, e2))) & (e2 = op(e1, e2))) | ~ sP81), inference(usedef, [], [e91])).
fof(e91, plain, (sP81 <=> (((e2 = unit) | ~ (e2 = op(e1, e2))) & (e2 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP81])])).
fof(f92, plain, ((((e2 = unit) | ~ (e3 = op(e1, e2))) & (e2 = op(e1, e3))) | ~ sP82), inference(usedef, [], [e92])).
fof(e92, plain, (sP82 <=> (((e2 = unit) | ~ (e3 = op(e1, e2))) & (e2 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP82])])).
fof(f93, plain, ((((e2 = unit) | ~ (e4 = op(e1, e2))) & (e2 = op(e1, e4))) | ~ sP83), inference(usedef, [], [e93])).
fof(e93, plain, (sP83 <=> (((e2 = unit) | ~ (e4 = op(e1, e2))) & (e2 = op(e1, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP83])])).
fof(f94, plain, ((((e2 = unit) | ~ (e0 = op(e2, e2))) & (e2 = op(e2, e0))) | ~ sP84), inference(usedef, [], [e94])).
fof(e94, plain, (sP84 <=> (((e2 = unit) | ~ (e0 = op(e2, e2))) & (e2 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP84])])).
fof(f95, plain, ((((e2 = unit) | ~ (e1 = op(e2, e2))) & (e2 = op(e2, e1))) | ~ sP85), inference(usedef, [], [e95])).
fof(e95, plain, (sP85 <=> (((e2 = unit) | ~ (e1 = op(e2, e2))) & (e2 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP85])])).
fof(f96, plain, ((((e2 = unit) | ~ (e2 = op(e2, e2))) & (e2 = op(e2, e2))) | ~ sP86), inference(usedef, [], [e96])).
fof(e96, plain, (sP86 <=> (((e2 = unit) | ~ (e2 = op(e2, e2))) & (e2 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP86])])).
fof(f97, plain, ((((e2 = unit) | ~ (e3 = op(e2, e2))) & (e2 = op(e2, e3))) | ~ sP87), inference(usedef, [], [e97])).
fof(e97, plain, (sP87 <=> (((e2 = unit) | ~ (e3 = op(e2, e2))) & (e2 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP87])])).
fof(f98, plain, ((((e2 = unit) | ~ (e4 = op(e2, e2))) & (e2 = op(e2, e4))) | ~ sP88), inference(usedef, [], [e98])).
fof(e98, plain, (sP88 <=> (((e2 = unit) | ~ (e4 = op(e2, e2))) & (e2 = op(e2, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP88])])).
fof(f99, plain, ((((e2 = unit) | ~ (e0 = op(e3, e2))) & (e2 = op(e3, e0))) | ~ sP89), inference(usedef, [], [e99])).
fof(e99, plain, (sP89 <=> (((e2 = unit) | ~ (e0 = op(e3, e2))) & (e2 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP89])])).
fof(f100, plain, ((((e2 = unit) | ~ (e1 = op(e3, e2))) & (e2 = op(e3, e1))) | ~ sP90), inference(usedef, [], [e100])).
fof(e100, plain, (sP90 <=> (((e2 = unit) | ~ (e1 = op(e3, e2))) & (e2 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP90])])).
fof(f101, plain, ((((e2 = unit) | ~ (e2 = op(e3, e2))) & (e2 = op(e3, e2))) | ~ sP91), inference(usedef, [], [e101])).
fof(e101, plain, (sP91 <=> (((e2 = unit) | ~ (e2 = op(e3, e2))) & (e2 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP91])])).
fof(f102, plain, ((((e2 = unit) | ~ (e3 = op(e3, e2))) & (e2 = op(e3, e3))) | ~ sP92), inference(usedef, [], [e102])).
fof(e102, plain, (sP92 <=> (((e2 = unit) | ~ (e3 = op(e3, e2))) & (e2 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP92])])).
fof(f103, plain, ((((e2 = unit) | ~ (e4 = op(e3, e2))) & (e2 = op(e3, e4))) | ~ sP93), inference(usedef, [], [e103])).
fof(e103, plain, (sP93 <=> (((e2 = unit) | ~ (e4 = op(e3, e2))) & (e2 = op(e3, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP93])])).
fof(f104, plain, ((((e2 = unit) | ~ (e0 = op(e4, e2))) & (e2 = op(e4, e0))) | ~ sP94), inference(usedef, [], [e104])).
fof(e104, plain, (sP94 <=> (((e2 = unit) | ~ (e0 = op(e4, e2))) & (e2 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP94])])).
fof(f105, plain, ((((e2 = unit) | ~ (e1 = op(e4, e2))) & (e2 = op(e4, e1))) | ~ sP95), inference(usedef, [], [e105])).
fof(e105, plain, (sP95 <=> (((e2 = unit) | ~ (e1 = op(e4, e2))) & (e2 = op(e4, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP95])])).
fof(f106, plain, ((((e3 = unit) | ~ (e0 = op(e0, e3))) & (op(e0, e0) = e3)) | ~ sP96), inference(usedef, [], [e106])).
fof(e106, plain, (sP96 <=> (((e3 = unit) | ~ (e0 = op(e0, e3))) & (op(e0, e0) = e3))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP96])])).
fof(f107, plain, ((((e3 = unit) | ~ (e1 = op(e0, e3))) & (e3 = op(e0, e1))) | ~ sP97), inference(usedef, [], [e107])).
fof(e107, plain, (sP97 <=> (((e3 = unit) | ~ (e1 = op(e0, e3))) & (e3 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP97])])).
fof(f108, plain, ((((e3 = unit) | ~ (e2 = op(e0, e3))) & (e3 = op(e0, e2))) | ~ sP98), inference(usedef, [], [e108])).
fof(e108, plain, (sP98 <=> (((e3 = unit) | ~ (e2 = op(e0, e3))) & (e3 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP98])])).
fof(f109, plain, ((((e3 = unit) | ~ (e3 = op(e0, e3))) & (e3 = op(e0, e3))) | ~ sP99), inference(usedef, [], [e109])).
fof(e109, plain, (sP99 <=> (((e3 = unit) | ~ (e3 = op(e0, e3))) & (e3 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP99])])).
fof(f110, plain, ((((e3 = unit) | ~ (e4 = op(e0, e3))) & (e3 = op(e0, e4))) | ~ sP100), inference(usedef, [], [e110])).
fof(e110, plain, (sP100 <=> (((e3 = unit) | ~ (e4 = op(e0, e3))) & (e3 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP100])])).
fof(f111, plain, ((((e3 = unit) | ~ (e0 = op(e1, e3))) & (e3 = op(e1, e0))) | ~ sP101), inference(usedef, [], [e111])).
fof(e111, plain, (sP101 <=> (((e3 = unit) | ~ (e0 = op(e1, e3))) & (e3 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP101])])).
fof(f112, plain, ((((e3 = unit) | ~ (e1 = op(e1, e3))) & (e3 = op(e1, e1))) | ~ sP102), inference(usedef, [], [e112])).
fof(e112, plain, (sP102 <=> (((e3 = unit) | ~ (e1 = op(e1, e3))) & (e3 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP102])])).
fof(f113, plain, ((((e3 = unit) | ~ (e2 = op(e1, e3))) & (e3 = op(e1, e2))) | ~ sP103), inference(usedef, [], [e113])).
fof(e113, plain, (sP103 <=> (((e3 = unit) | ~ (e2 = op(e1, e3))) & (e3 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP103])])).
fof(f114, plain, ((((e3 = unit) | ~ (e3 = op(e1, e3))) & (e3 = op(e1, e3))) | ~ sP104), inference(usedef, [], [e114])).
fof(e114, plain, (sP104 <=> (((e3 = unit) | ~ (e3 = op(e1, e3))) & (e3 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP104])])).
fof(f115, plain, ((((e3 = unit) | ~ (e4 = op(e1, e3))) & (e3 = op(e1, e4))) | ~ sP105), inference(usedef, [], [e115])).
fof(e115, plain, (sP105 <=> (((e3 = unit) | ~ (e4 = op(e1, e3))) & (e3 = op(e1, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP105])])).
fof(f116, plain, ((((e3 = unit) | ~ (e0 = op(e2, e3))) & (e3 = op(e2, e0))) | ~ sP106), inference(usedef, [], [e116])).
fof(e116, plain, (sP106 <=> (((e3 = unit) | ~ (e0 = op(e2, e3))) & (e3 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP106])])).
fof(f117, plain, ((((e3 = unit) | ~ (e1 = op(e2, e3))) & (e3 = op(e2, e1))) | ~ sP107), inference(usedef, [], [e117])).
fof(e117, plain, (sP107 <=> (((e3 = unit) | ~ (e1 = op(e2, e3))) & (e3 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP107])])).
fof(f118, plain, ((((e3 = unit) | ~ (e2 = op(e2, e3))) & (e3 = op(e2, e2))) | ~ sP108), inference(usedef, [], [e118])).
fof(e118, plain, (sP108 <=> (((e3 = unit) | ~ (e2 = op(e2, e3))) & (e3 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP108])])).
fof(f119, plain, ((((e3 = unit) | ~ (e3 = op(e2, e3))) & (e3 = op(e2, e3))) | ~ sP109), inference(usedef, [], [e119])).
fof(e119, plain, (sP109 <=> (((e3 = unit) | ~ (e3 = op(e2, e3))) & (e3 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP109])])).
fof(f120, plain, ((((e3 = unit) | ~ (e4 = op(e2, e3))) & (e3 = op(e2, e4))) | ~ sP110), inference(usedef, [], [e120])).
fof(e120, plain, (sP110 <=> (((e3 = unit) | ~ (e4 = op(e2, e3))) & (e3 = op(e2, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP110])])).
fof(f121, plain, ((((e3 = unit) | ~ (e0 = op(e3, e3))) & (e3 = op(e3, e0))) | ~ sP111), inference(usedef, [], [e121])).
fof(e121, plain, (sP111 <=> (((e3 = unit) | ~ (e0 = op(e3, e3))) & (e3 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP111])])).
fof(f122, plain, ((((e3 = unit) | ~ (e1 = op(e3, e3))) & (e3 = op(e3, e1))) | ~ sP112), inference(usedef, [], [e122])).
fof(e122, plain, (sP112 <=> (((e3 = unit) | ~ (e1 = op(e3, e3))) & (e3 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP112])])).
fof(f123, plain, ((((e3 = unit) | ~ (e2 = op(e3, e3))) & (e3 = op(e3, e2))) | ~ sP113), inference(usedef, [], [e123])).
fof(e123, plain, (sP113 <=> (((e3 = unit) | ~ (e2 = op(e3, e3))) & (e3 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP113])])).
fof(f124, plain, ((((e3 = unit) | ~ (e3 = op(e3, e3))) & (e3 = op(e3, e3))) | ~ sP114), inference(usedef, [], [e124])).
fof(e124, plain, (sP114 <=> (((e3 = unit) | ~ (e3 = op(e3, e3))) & (e3 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP114])])).
fof(f125, plain, ((((e3 = unit) | ~ (e4 = op(e3, e3))) & (e3 = op(e3, e4))) | ~ sP115), inference(usedef, [], [e125])).
fof(e125, plain, (sP115 <=> (((e3 = unit) | ~ (e4 = op(e3, e3))) & (e3 = op(e3, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP115])])).
fof(f126, plain, ((((e3 = unit) | ~ (e0 = op(e4, e3))) & (e3 = op(e4, e0))) | ~ sP116), inference(usedef, [], [e126])).
fof(e126, plain, (sP116 <=> (((e3 = unit) | ~ (e0 = op(e4, e3))) & (e3 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP116])])).
fof(f127, plain, ((((e3 = unit) | ~ (e1 = op(e4, e3))) & (e3 = op(e4, e1))) | ~ sP117), inference(usedef, [], [e127])).
fof(e127, plain, (sP117 <=> (((e3 = unit) | ~ (e1 = op(e4, e3))) & (e3 = op(e4, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP117])])).
fof(f128, plain, ((((e4 = unit) | ~ (e0 = op(e0, e4))) & (op(e0, e0) = e4)) | ~ sP118), inference(usedef, [], [e128])).
fof(e128, plain, (sP118 <=> (((e4 = unit) | ~ (e0 = op(e0, e4))) & (op(e0, e0) = e4))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP118])])).
fof(f129, plain, ((((e4 = unit) | ~ (e1 = op(e0, e4))) & (e4 = op(e0, e1))) | ~ sP119), inference(usedef, [], [e129])).
fof(e129, plain, (sP119 <=> (((e4 = unit) | ~ (e1 = op(e0, e4))) & (e4 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP119])])).
fof(f130, plain, ((((e4 = unit) | ~ (e2 = op(e0, e4))) & (e4 = op(e0, e2))) | ~ sP120), inference(usedef, [], [e130])).
fof(e130, plain, (sP120 <=> (((e4 = unit) | ~ (e2 = op(e0, e4))) & (e4 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP120])])).
fof(f131, plain, ((((e4 = unit) | ~ (e3 = op(e0, e4))) & (e4 = op(e0, e3))) | ~ sP121), inference(usedef, [], [e131])).
fof(e131, plain, (sP121 <=> (((e4 = unit) | ~ (e3 = op(e0, e4))) & (e4 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP121])])).
fof(f132, plain, ((((e4 = unit) | ~ (e4 = op(e0, e4))) & (e4 = op(e0, e4))) | ~ sP122), inference(usedef, [], [e132])).
fof(e132, plain, (sP122 <=> (((e4 = unit) | ~ (e4 = op(e0, e4))) & (e4 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP122])])).
fof(f133, plain, ((((e4 = unit) | ~ (e0 = op(e1, e4))) & (e4 = op(e1, e0))) | ~ sP123), inference(usedef, [], [e133])).
fof(e133, plain, (sP123 <=> (((e4 = unit) | ~ (e0 = op(e1, e4))) & (e4 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP123])])).
fof(f134, plain, ((((e4 = unit) | ~ (e1 = op(e1, e4))) & (e4 = op(e1, e1))) | ~ sP124), inference(usedef, [], [e134])).
fof(e134, plain, (sP124 <=> (((e4 = unit) | ~ (e1 = op(e1, e4))) & (e4 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP124])])).
fof(f135, plain, ((((e4 = unit) | ~ (e2 = op(e1, e4))) & (e4 = op(e1, e2))) | ~ sP125), inference(usedef, [], [e135])).
fof(e135, plain, (sP125 <=> (((e4 = unit) | ~ (e2 = op(e1, e4))) & (e4 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP125])])).
fof(f136, plain, ((((e4 = unit) | ~ (e3 = op(e1, e4))) & (e4 = op(e1, e3))) | ~ sP126), inference(usedef, [], [e136])).
fof(e136, plain, (sP126 <=> (((e4 = unit) | ~ (e3 = op(e1, e4))) & (e4 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP126])])).
fof(f137, plain, ((((e4 = unit) | ~ (e4 = op(e1, e4))) & (e4 = op(e1, e4))) | ~ sP127), inference(usedef, [], [e137])).
fof(e137, plain, (sP127 <=> (((e4 = unit) | ~ (e4 = op(e1, e4))) & (e4 = op(e1, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP127])])).
fof(f138, plain, ((((e4 = unit) | ~ (e0 = op(e2, e4))) & (e4 = op(e2, e0))) | ~ sP128), inference(usedef, [], [e138])).
fof(e138, plain, (sP128 <=> (((e4 = unit) | ~ (e0 = op(e2, e4))) & (e4 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP128])])).
fof(f139, plain, ((((e4 = unit) | ~ (e1 = op(e2, e4))) & (e4 = op(e2, e1))) | ~ sP129), inference(usedef, [], [e139])).
fof(e139, plain, (sP129 <=> (((e4 = unit) | ~ (e1 = op(e2, e4))) & (e4 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP129])])).
fof(f140, plain, ((((e4 = unit) | ~ (e2 = op(e2, e4))) & (e4 = op(e2, e2))) | ~ sP130), inference(usedef, [], [e140])).
fof(e140, plain, (sP130 <=> (((e4 = unit) | ~ (e2 = op(e2, e4))) & (e4 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP130])])).
fof(f141, plain, ((((e4 = unit) | ~ (e3 = op(e2, e4))) & (e4 = op(e2, e3))) | ~ sP131), inference(usedef, [], [e141])).
fof(e141, plain, (sP131 <=> (((e4 = unit) | ~ (e3 = op(e2, e4))) & (e4 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP131])])).
fof(f142, plain, ((((e4 = unit) | ~ (e4 = op(e2, e4))) & (e4 = op(e2, e4))) | ~ sP132), inference(usedef, [], [e142])).
fof(e142, plain, (sP132 <=> (((e4 = unit) | ~ (e4 = op(e2, e4))) & (e4 = op(e2, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP132])])).
fof(f143, plain, ((((e4 = unit) | ~ (e0 = op(e3, e4))) & (e4 = op(e3, e0))) | ~ sP133), inference(usedef, [], [e143])).
fof(e143, plain, (sP133 <=> (((e4 = unit) | ~ (e0 = op(e3, e4))) & (e4 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP133])])).
fof(f144, plain, ((((e4 = unit) | ~ (e1 = op(e3, e4))) & (e4 = op(e3, e1))) | ~ sP134), inference(usedef, [], [e144])).
fof(e144, plain, (sP134 <=> (((e4 = unit) | ~ (e1 = op(e3, e4))) & (e4 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP134])])).
fof(f145, plain, ((((e4 = unit) | ~ (e2 = op(e3, e4))) & (e4 = op(e3, e2))) | ~ sP135), inference(usedef, [], [e145])).
fof(e145, plain, (sP135 <=> (((e4 = unit) | ~ (e2 = op(e3, e4))) & (e4 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP135])])).
fof(f146, plain, ((((e4 = unit) | ~ (e3 = op(e3, e4))) & (e4 = op(e3, e3))) | ~ sP136), inference(usedef, [], [e146])).
fof(e146, plain, (sP136 <=> (((e4 = unit) | ~ (e3 = op(e3, e4))) & (e4 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP136])])).
fof(f147, plain, ((((e4 = unit) | ~ (e4 = op(e3, e4))) & (e4 = op(e3, e4))) | ~ sP137), inference(usedef, [], [e147])).
fof(e147, plain, (sP137 <=> (((e4 = unit) | ~ (e4 = op(e3, e4))) & (e4 = op(e3, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP137])])).
fof(f148, plain, ((((e4 = unit) | ~ (e0 = op(e4, e4))) & (e4 = op(e4, e0))) | ~ sP138), inference(usedef, [], [e148])).
fof(e148, plain, (sP138 <=> (((e4 = unit) | ~ (e0 = op(e4, e4))) & (e4 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP138])])).
fof(f149, plain, ((((e4 = unit) | ~ (e1 = op(e4, e4))) & (e4 = op(e4, e1))) | ~ sP139), inference(usedef, [], [e149])).
fof(e149, plain, (sP139 <=> (((e4 = unit) | ~ (e1 = op(e4, e4))) & (e4 = op(e4, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP139])])).
fof(f9, plain, (((e4 = op(op(e4, e4), e4)) | ~ (e4 = op(op(e4, e4), e4)) | (op(e4, e4) = op(e4, e4))) & ((e3 = op(op(e4, e3), e4)) | ~ (e4 = op(op(e4, e3), e3)) | (op(e3, e4) = op(e4, e3))) & ((e2 = op(op(e4, e2), e4)) | ~ (e4 = op(op(e4, e2), e2)) | (op(e2, e4) = op(e4, e2))) & ((e1 = op(op(e4, e1), e4)) | ~ (e4 = op(op(e4, e1), e1)) | (op(e1, e4) = op(e4, e1))) & ((e0 = op(op(e4, e0), e4)) | ~ (e4 = op(op(e4, e0), e0)) | (op(e0, e4) = op(e4, e0))) & ((e4 = op(op(e3, e4), e3)) | ~ (e3 = op(op(e3, e4), e4)) | (op(e3, e4) = op(e4, e3))) & ((e3 = op(op(e3, e3), e3)) | ~ (e3 = op(op(e3, e3), e3)) | (op(e3, e3) = op(e3, e3))) & ((e2 = op(op(e3, e2), e3)) | ~ (e3 = op(op(e3, e2), e2)) | (op(e2, e3) = op(e3, e2))) & ((e1 = op(op(e3, e1), e3)) | ~ (e3 = op(op(e3, e1), e1)) | (op(e1, e3) = op(e3, e1))) & ((e0 = op(op(e3, e0), e3)) | ~ (e3 = op(op(e3, e0), e0)) | (op(e0, e3) = op(e3, e0))) & ((e4 = op(op(e2, e4), e2)) | ~ (e2 = op(op(e2, e4), e4)) | (op(e2, e4) = op(e4, e2))) & ((e3 = op(op(e2, e3), e2)) | ~ (e2 = op(op(e2, e3), e3)) | (op(e2, e3) = op(e3, e2))) & ((e2 = op(op(e2, e2), e2)) | ~ (e2 = op(op(e2, e2), e2)) | (op(e2, e2) = op(e2, e2))) & ((e1 = op(op(e2, e1), e2)) | ~ (e2 = op(op(e2, e1), e1)) | (op(e1, e2) = op(e2, e1))) & ((e0 = op(op(e2, e0), e2)) | ~ (e2 = op(op(e2, e0), e0)) | (op(e0, e2) = op(e2, e0))) & ((e4 = op(op(e1, e4), e1)) | ~ (e1 = op(op(e1, e4), e4)) | (op(e1, e4) = op(e4, e1))) & ((e3 = op(op(e1, e3), e1)) | ~ (e1 = op(op(e1, e3), e3)) | (op(e1, e3) = op(e3, e1))) & ((e2 = op(op(e1, e2), e1)) | ~ (e1 = op(op(e1, e2), e2)) | (op(e1, e2) = op(e2, e1))) & ((e1 = op(op(e1, e1), e1)) | ~ (e1 = op(op(e1, e1), e1)) | (op(e1, e1) = op(e1, e1))) & ((e0 = op(op(e1, e0), e1)) | ~ (e1 = op(op(e1, e0), e0)) | (op(e0, e1) = op(e1, e0))) & ((e4 = op(op(e0, e4), e0)) | ~ (e0 = op(op(e0, e4), e4)) | (op(e0, e4) = op(e4, e0))) & ((e3 = op(op(e0, e3), e0)) | ~ (e0 = op(op(e0, e3), e3)) | (op(e0, e3) = op(e3, e0))) & ((e2 = op(op(e0, e2), e0)) | ~ (e0 = op(op(e0, e2), e2)) | (op(e0, e2) = op(e2, e0))) & ((e1 = op(op(e0, e1), e0)) | ~ (e0 = op(op(e0, e1), e1)) | (op(e0, e1) = op(e1, e0))) & ((e0 = op(op(e0, e0), e0)) | ~ (e0 = op(op(e0, e0), e0)) | (op(e0, e0) = op(e0, e0))) & ((e4 = op(e4, e4)) | ~ (e4 = op(e4, e4)) | ~ (e4 = op(e4, e4))) & ((e4 = op(e4, e3)) | ~ (e4 = op(e3, e3)) | ~ (e3 = op(e4, e4))) & ((e4 = op(e4, e2)) | ~ (e4 = op(e2, e2)) | ~ (e2 = op(e4, e4))) & ((e4 = op(e4, e1)) | ~ (e4 = op(e1, e1)) | ~ (e1 = op(e4, e4))) & ((e4 = op(e4, e0)) | ~ (op(e0, e0) = e4) | ~ (e0 = op(e4, e4))) & ((e3 = op(e3, e4)) | ~ (e3 = op(e4, e4)) | ~ (e4 = op(e3, e3))) & ((e3 = op(e3, e3)) | ~ (e3 = op(e3, e3)) | ~ (e3 = op(e3, e3))) & ((e3 = op(e3, e2)) | ~ (e3 = op(e2, e2)) | ~ (e2 = op(e3, e3))) & ((e3 = op(e3, e1)) | ~ (e3 = op(e1, e1)) | ~ (e1 = op(e3, e3))) & ((e3 = op(e3, e0)) | ~ (op(e0, e0) = e3) | ~ (e0 = op(e3, e3))) & ((e2 = op(e2, e4)) | ~ (e2 = op(e4, e4)) | ~ (e4 = op(e2, e2))) & ((e2 = op(e2, e3)) | ~ (e2 = op(e3, e3)) | ~ (e3 = op(e2, e2))) & ((e2 = op(e2, e2)) | ~ (e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & ((e2 = op(e2, e1)) | ~ (e2 = op(e1, e1)) | ~ (e1 = op(e2, e2))) & ((e2 = op(e2, e0)) | ~ (op(e0, e0) = e2) | ~ (e0 = op(e2, e2))) & ((e1 = op(e1, e4)) | ~ (e1 = op(e4, e4)) | ~ (e4 = op(e1, e1))) & ((e1 = op(e1, e3)) | ~ (e1 = op(e3, e3)) | ~ (e3 = op(e1, e1))) & ((e1 = op(e1, e2)) | ~ (e1 = op(e2, e2)) | ~ (e2 = op(e1, e1))) & ((e1 = op(e1, e1)) | ~ (e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & ((e1 = op(e1, e0)) | ~ (op(e0, e0) = e1) | ~ (e0 = op(e1, e1))) & ((e0 = op(e0, e4)) | ~ (e0 = op(e4, e4)) | ~ (op(e0, e0) = e4)) & ((e0 = op(e0, e3)) | ~ (e0 = op(e3, e3)) | ~ (op(e0, e0) = e3)) & ((e0 = op(e0, e2)) | ~ (e0 = op(e2, e2)) | ~ (op(e0, e0) = e2)) & ((e0 = op(e0, e1)) | ~ (e0 = op(e1, e1)) | ~ (op(e0, e0) = e1)) & ((e0 = op(e0, e0)) | ~ (e0 = op(e0, e0)) | ~ (e0 = op(e0, e0))) & ((((e4 = unit) | ~ (e4 = op(e4, e4))) & (e4 = op(e4, e4))) | (((e4 = unit) | ~ (e3 = op(e4, e4))) & (e4 = op(e4, e3))) | (((e4 = unit) | ~ (e2 = op(e4, e4))) & (e4 = op(e4, e2))) | (((e4 = unit) | ~ (e1 = op(e4, e4))) & (e4 = op(e4, e1))) | (((e4 = unit) | ~ (e0 = op(e4, e4))) & (e4 = op(e4, e0))) | (((e4 = unit) | ~ (e4 = op(e3, e4))) & (e4 = op(e3, e4))) | (((e4 = unit) | ~ (e3 = op(e3, e4))) & (e4 = op(e3, e3))) | (((e4 = unit) | ~ (e2 = op(e3, e4))) & (e4 = op(e3, e2))) | (((e4 = unit) | ~ (e1 = op(e3, e4))) & (e4 = op(e3, e1))) | (((e4 = unit) | ~ (e0 = op(e3, e4))) & (e4 = op(e3, e0))) | (((e4 = unit) | ~ (e4 = op(e2, e4))) & (e4 = op(e2, e4))) | (((e4 = unit) | ~ (e3 = op(e2, e4))) & (e4 = op(e2, e3))) | (((e4 = unit) | ~ (e2 = op(e2, e4))) & (e4 = op(e2, e2))) | (((e4 = unit) | ~ (e1 = op(e2, e4))) & (e4 = op(e2, e1))) | (((e4 = unit) | ~ (e0 = op(e2, e4))) & (e4 = op(e2, e0))) | (((e4 = unit) | ~ (e4 = op(e1, e4))) & (e4 = op(e1, e4))) | (((e4 = unit) | ~ (e3 = op(e1, e4))) & (e4 = op(e1, e3))) | (((e4 = unit) | ~ (e2 = op(e1, e4))) & (e4 = op(e1, e2))) | (((e4 = unit) | ~ (e1 = op(e1, e4))) & (e4 = op(e1, e1))) | (((e4 = unit) | ~ (e0 = op(e1, e4))) & (e4 = op(e1, e0))) | (((e4 = unit) | ~ (e4 = op(e0, e4))) & (e4 = op(e0, e4))) | (((e4 = unit) | ~ (e3 = op(e0, e4))) & (e4 = op(e0, e3))) | (((e4 = unit) | ~ (e2 = op(e0, e4))) & (e4 = op(e0, e2))) | (((e4 = unit) | ~ (e1 = op(e0, e4))) & (e4 = op(e0, e1))) | (((e4 = unit) | ~ (e0 = op(e0, e4))) & (op(e0, e0) = e4))) & ((((e3 = unit) | ~ (e4 = op(e4, e3))) & (e3 = op(e4, e4))) | (((e3 = unit) | ~ (e3 = op(e4, e3))) & (e3 = op(e4, e3))) | (((e3 = unit) | ~ (e2 = op(e4, e3))) & (e3 = op(e4, e2))) | (((e3 = unit) | ~ (e1 = op(e4, e3))) & (e3 = op(e4, e1))) | (((e3 = unit) | ~ (e0 = op(e4, e3))) & (e3 = op(e4, e0))) | (((e3 = unit) | ~ (e4 = op(e3, e3))) & (e3 = op(e3, e4))) | (((e3 = unit) | ~ (e3 = op(e3, e3))) & (e3 = op(e3, e3))) | (((e3 = unit) | ~ (e2 = op(e3, e3))) & (e3 = op(e3, e2))) | (((e3 = unit) | ~ (e1 = op(e3, e3))) & (e3 = op(e3, e1))) | (((e3 = unit) | ~ (e0 = op(e3, e3))) & (e3 = op(e3, e0))) | (((e3 = unit) | ~ (e4 = op(e2, e3))) & (e3 = op(e2, e4))) | (((e3 = unit) | ~ (e3 = op(e2, e3))) & (e3 = op(e2, e3))) | (((e3 = unit) | ~ (e2 = op(e2, e3))) & (e3 = op(e2, e2))) | (((e3 = unit) | ~ (e1 = op(e2, e3))) & (e3 = op(e2, e1))) | (((e3 = unit) | ~ (e0 = op(e2, e3))) & (e3 = op(e2, e0))) | (((e3 = unit) | ~ (e4 = op(e1, e3))) & (e3 = op(e1, e4))) | (((e3 = unit) | ~ (e3 = op(e1, e3))) & (e3 = op(e1, e3))) | (((e3 = unit) | ~ (e2 = op(e1, e3))) & (e3 = op(e1, e2))) | (((e3 = unit) | ~ (e1 = op(e1, e3))) & (e3 = op(e1, e1))) | (((e3 = unit) | ~ (e0 = op(e1, e3))) & (e3 = op(e1, e0))) | (((e3 = unit) | ~ (e4 = op(e0, e3))) & (e3 = op(e0, e4))) | (((e3 = unit) | ~ (e3 = op(e0, e3))) & (e3 = op(e0, e3))) | (((e3 = unit) | ~ (e2 = op(e0, e3))) & (e3 = op(e0, e2))) | (((e3 = unit) | ~ (e1 = op(e0, e3))) & (e3 = op(e0, e1))) | (((e3 = unit) | ~ (e0 = op(e0, e3))) & (op(e0, e0) = e3))) & ((((e2 = unit) | ~ (e4 = op(e4, e2))) & (e2 = op(e4, e4))) | (((e2 = unit) | ~ (e3 = op(e4, e2))) & (e2 = op(e4, e3))) | (((e2 = unit) | ~ (e2 = op(e4, e2))) & (e2 = op(e4, e2))) | (((e2 = unit) | ~ (e1 = op(e4, e2))) & (e2 = op(e4, e1))) | (((e2 = unit) | ~ (e0 = op(e4, e2))) & (e2 = op(e4, e0))) | (((e2 = unit) | ~ (e4 = op(e3, e2))) & (e2 = op(e3, e4))) | (((e2 = unit) | ~ (e3 = op(e3, e2))) & (e2 = op(e3, e3))) | (((e2 = unit) | ~ (e2 = op(e3, e2))) & (e2 = op(e3, e2))) | (((e2 = unit) | ~ (e1 = op(e3, e2))) & (e2 = op(e3, e1))) | (((e2 = unit) | ~ (e0 = op(e3, e2))) & (e2 = op(e3, e0))) | (((e2 = unit) | ~ (e4 = op(e2, e2))) & (e2 = op(e2, e4))) | (((e2 = unit) | ~ (e3 = op(e2, e2))) & (e2 = op(e2, e3))) | (((e2 = unit) | ~ (e2 = op(e2, e2))) & (e2 = op(e2, e2))) | (((e2 = unit) | ~ (e1 = op(e2, e2))) & (e2 = op(e2, e1))) | (((e2 = unit) | ~ (e0 = op(e2, e2))) & (e2 = op(e2, e0))) | (((e2 = unit) | ~ (e4 = op(e1, e2))) & (e2 = op(e1, e4))) | (((e2 = unit) | ~ (e3 = op(e1, e2))) & (e2 = op(e1, e3))) | (((e2 = unit) | ~ (e2 = op(e1, e2))) & (e2 = op(e1, e2))) | (((e2 = unit) | ~ (e1 = op(e1, e2))) & (e2 = op(e1, e1))) | (((e2 = unit) | ~ (e0 = op(e1, e2))) & (e2 = op(e1, e0))) | (((e2 = unit) | ~ (e4 = op(e0, e2))) & (e2 = op(e0, e4))) | (((e2 = unit) | ~ (e3 = op(e0, e2))) & (e2 = op(e0, e3))) | (((e2 = unit) | ~ (e2 = op(e0, e2))) & (e2 = op(e0, e2))) | (((e2 = unit) | ~ (e1 = op(e0, e2))) & (e2 = op(e0, e1))) | (((e2 = unit) | ~ (e0 = op(e0, e2))) & (op(e0, e0) = e2))) & ((((e1 = unit) | ~ (e4 = op(e4, e1))) & (e1 = op(e4, e4))) | (((e1 = unit) | ~ (e3 = op(e4, e1))) & (e1 = op(e4, e3))) | (((e1 = unit) | ~ (e2 = op(e4, e1))) & (e1 = op(e4, e2))) | (((e1 = unit) | ~ (e1 = op(e4, e1))) & (e1 = op(e4, e1))) | (((e1 = unit) | ~ (e0 = op(e4, e1))) & (e1 = op(e4, e0))) | (((e1 = unit) | ~ (e4 = op(e3, e1))) & (e1 = op(e3, e4))) | (((e1 = unit) | ~ (e3 = op(e3, e1))) & (e1 = op(e3, e3))) | (((e1 = unit) | ~ (e2 = op(e3, e1))) & (e1 = op(e3, e2))) | (((e1 = unit) | ~ (e1 = op(e3, e1))) & (e1 = op(e3, e1))) | (((e1 = unit) | ~ (e0 = op(e3, e1))) & (e1 = op(e3, e0))) | (((e1 = unit) | ~ (e4 = op(e2, e1))) & (e1 = op(e2, e4))) | (((e1 = unit) | ~ (e3 = op(e2, e1))) & (e1 = op(e2, e3))) | (((e1 = unit) | ~ (e2 = op(e2, e1))) & (e1 = op(e2, e2))) | (((e1 = unit) | ~ (e1 = op(e2, e1))) & (e1 = op(e2, e1))) | (((e1 = unit) | ~ (e0 = op(e2, e1))) & (e1 = op(e2, e0))) | (((e1 = unit) | ~ (e4 = op(e1, e1))) & (e1 = op(e1, e4))) | (((e1 = unit) | ~ (e3 = op(e1, e1))) & (e1 = op(e1, e3))) | (((e1 = unit) | ~ (e2 = op(e1, e1))) & (e1 = op(e1, e2))) | (((e1 = unit) | ~ (e1 = op(e1, e1))) & (e1 = op(e1, e1))) | (((e1 = unit) | ~ (e0 = op(e1, e1))) & (e1 = op(e1, e0))) | (((e1 = unit) | ~ (e4 = op(e0, e1))) & (e1 = op(e0, e4))) | (((e1 = unit) | ~ (e3 = op(e0, e1))) & (e1 = op(e0, e3))) | (((e1 = unit) | ~ (e2 = op(e0, e1))) & (e1 = op(e0, e2))) | (((e1 = unit) | ~ (e1 = op(e0, e1))) & (e1 = op(e0, e1))) | (((e1 = unit) | ~ (e0 = op(e0, e1))) & (op(e0, e0) = e1))) & ((((e0 = unit) | ~ (e4 = op(e4, e0))) & (e0 = op(e4, e4))) | (((e0 = unit) | ~ (e3 = op(e4, e0))) & (e0 = op(e4, e3))) | (((e0 = unit) | ~ (e2 = op(e4, e0))) & (e0 = op(e4, e2))) | (((e0 = unit) | ~ (e1 = op(e4, e0))) & (e0 = op(e4, e1))) | (((e0 = unit) | ~ (e0 = op(e4, e0))) & (e0 = op(e4, e0))) | (((e0 = unit) | ~ (e4 = op(e3, e0))) & (e0 = op(e3, e4))) | (((e0 = unit) | ~ (e3 = op(e3, e0))) & (e0 = op(e3, e3))) | (((e0 = unit) | ~ (e2 = op(e3, e0))) & (e0 = op(e3, e2))) | (((e0 = unit) | ~ (e1 = op(e3, e0))) & (e0 = op(e3, e1))) | (((e0 = unit) | ~ (e0 = op(e3, e0))) & (e0 = op(e3, e0))) | (((e0 = unit) | ~ (e4 = op(e2, e0))) & (e0 = op(e2, e4))) | (((e0 = unit) | ~ (e3 = op(e2, e0))) & (e0 = op(e2, e3))) | (((e0 = unit) | ~ (e2 = op(e2, e0))) & (e0 = op(e2, e2))) | (((e0 = unit) | ~ (e1 = op(e2, e0))) & (e0 = op(e2, e1))) | (((e0 = unit) | ~ (e0 = op(e2, e0))) & (e0 = op(e2, e0))) | (((e0 = unit) | ~ (e4 = op(e1, e0))) & (e0 = op(e1, e4))) | (((e0 = unit) | ~ (e3 = op(e1, e0))) & (e0 = op(e1, e3))) | (((e0 = unit) | ~ (e2 = op(e1, e0))) & (e0 = op(e1, e2))) | (((e0 = unit) | ~ (e1 = op(e1, e0))) & (e0 = op(e1, e1))) | (((e0 = unit) | ~ (e0 = op(e1, e0))) & (e0 = op(e1, e0))) | (((e0 = unit) | ~ (op(e0, e0) = e4)) & (e0 = op(e0, e4))) | (((e0 = unit) | ~ (op(e0, e0) = e3)) & (e0 = op(e0, e3))) | (((e0 = unit) | ~ (op(e0, e0) = e2)) & (e0 = op(e0, e2))) | (((e0 = unit) | ~ (op(e0, e0) = e1)) & (e0 = op(e0, e1))) | (((e0 = unit) | ~ (e0 = op(e0, e0))) & (e0 = op(e0, e0)))) & (((~ (e4 = op(e4, op(e4, e4))) | (e4 = op(e4, op(e4, e4)))) & (~ (e3 = op(e4, op(e3, e4))) | (e4 = op(e3, op(e3, e4)))) & (~ (e2 = op(e4, op(e2, e4))) | (e4 = op(e2, op(e2, e4)))) & (~ (e1 = op(e4, op(e1, e4))) | (e4 = op(e1, op(e1, e4)))) & (~ (e0 = op(e4, op(e0, e4))) | (e4 = op(e0, op(e0, e4))))) | ((~ (e4 = op(e3, op(e4, e3))) | (e3 = op(e4, op(e4, e3)))) & (~ (e3 = op(e3, op(e3, e3))) | (e3 = op(e3, op(e3, e3)))) & (~ (e2 = op(e3, op(e2, e3))) | (e3 = op(e2, op(e2, e3)))) & (~ (e1 = op(e3, op(e1, e3))) | (e3 = op(e1, op(e1, e3)))) & (~ (e0 = op(e3, op(e0, e3))) | (e3 = op(e0, op(e0, e3))))) | ((~ (e4 = op(e2, op(e4, e2))) | (e2 = op(e4, op(e4, e2)))) & (~ (e3 = op(e2, op(e3, e2))) | (e2 = op(e3, op(e3, e2)))) & (~ (e2 = op(e2, op(e2, e2))) | (e2 = op(e2, op(e2, e2)))) & (~ (e1 = op(e2, op(e1, e2))) | (e2 = op(e1, op(e1, e2)))) & (~ (e0 = op(e2, op(e0, e2))) | (e2 = op(e0, op(e0, e2))))) | ((~ (e4 = op(e1, op(e4, e1))) | (e1 = op(e4, op(e4, e1)))) & (~ (e3 = op(e1, op(e3, e1))) | (e1 = op(e3, op(e3, e1)))) & (~ (e2 = op(e1, op(e2, e1))) | (e1 = op(e2, op(e2, e1)))) & (~ (e1 = op(e1, op(e1, e1))) | (e1 = op(e1, op(e1, e1)))) & (~ (e0 = op(e1, op(e0, e1))) | (e1 = op(e0, op(e0, e1))))) | ((~ (e4 = op(e0, op(e4, e0))) | (e0 = op(e4, op(e4, e0)))) & (~ (e3 = op(e0, op(e3, e0))) | (e0 = op(e3, op(e3, e0)))) & (~ (e2 = op(e0, op(e2, e0))) | (e0 = op(e2, op(e2, e0)))) & (~ (e1 = op(e0, op(e1, e0))) | (e0 = op(e1, op(e1, e0)))) & (~ (e0 = op(e0, op(e0, e0))) | (e0 = op(e0, op(e0, e0)))))) & ((~ (e4 = op(e4, e4)) & (e4 = op(e4, e4))) | (~ (e4 = op(e4, e3)) & (e3 = op(e4, e4))) | (~ (e4 = op(e4, e2)) & (e2 = op(e4, e4))) | (~ (e4 = op(e4, e1)) & (e1 = op(e4, e4))) | (~ (e4 = op(e4, e0)) & (e0 = op(e4, e4))) | (~ (e3 = op(e3, e4)) & (e4 = op(e3, e3))) | (~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | (~ (e3 = op(e3, e2)) & (e2 = op(e3, e3))) | (~ (e3 = op(e3, e1)) & (e1 = op(e3, e3))) | (~ (e3 = op(e3, e0)) & (e0 = op(e3, e3))) | (~ (e2 = op(e2, e4)) & (e4 = op(e2, e2))) | (~ (e2 = op(e2, e3)) & (e3 = op(e2, e2))) | (~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | (~ (e2 = op(e2, e1)) & (e1 = op(e2, e2))) | (~ (e2 = op(e2, e0)) & (e0 = op(e2, e2))) | (~ (e1 = op(e1, e4)) & (e4 = op(e1, e1))) | (~ (e1 = op(e1, e3)) & (e3 = op(e1, e1))) | (~ (e1 = op(e1, e2)) & (e2 = op(e1, e1))) | (~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | (~ (e1 = op(e1, e0)) & (e0 = op(e1, e1))) | (~ (e0 = op(e0, e4)) & (op(e0, e0) = e4)) | (~ (e0 = op(e0, e3)) & (op(e0, e0) = e3)) | (~ (e0 = op(e0, e2)) & (op(e0, e0) = e2)) | (~ (e0 = op(e0, e1)) & (op(e0, e0) = e1)) | (~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)))) & ((~ (e4 = op(e4, e4)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e1, e1)) & ~ (op(e0, e0) = e4)) | (~ (e3 = op(e4, e4)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e1, e1)) & ~ (op(e0, e0) = e3)) | (~ (e2 = op(e4, e4)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e1, e1)) & ~ (op(e0, e0) = e2)) | (~ (e1 = op(e4, e4)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e1, e1)) & ~ (op(e0, e0) = e1)) | (~ (e0 = op(e4, e4)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e0, e0))))), inference(ennf_transformation, [], [f8])).
fof(f8, plain, ~ ((~ (e4 = op(op(e4, e4), e4)) & (e4 = op(op(e4, e4), e4)) & ~ (op(e4, e4) = op(e4, e4))) | (~ (e3 = op(op(e4, e3), e4)) & (e4 = op(op(e4, e3), e3)) & ~ (op(e3, e4) = op(e4, e3))) | (~ (e2 = op(op(e4, e2), e4)) & (e4 = op(op(e4, e2), e2)) & ~ (op(e2, e4) = op(e4, e2))) | (~ (e1 = op(op(e4, e1), e4)) & (e4 = op(op(e4, e1), e1)) & ~ (op(e1, e4) = op(e4, e1))) | (~ (e0 = op(op(e4, e0), e4)) & (e4 = op(op(e4, e0), e0)) & ~ (op(e0, e4) = op(e4, e0))) | (~ (e4 = op(op(e3, e4), e3)) & (e3 = op(op(e3, e4), e4)) & ~ (op(e3, e4) = op(e4, e3))) | (~ (e3 = op(op(e3, e3), e3)) & (e3 = op(op(e3, e3), e3)) & ~ (op(e3, e3) = op(e3, e3))) | (~ (e2 = op(op(e3, e2), e3)) & (e3 = op(op(e3, e2), e2)) & ~ (op(e2, e3) = op(e3, e2))) | (~ (e1 = op(op(e3, e1), e3)) & (e3 = op(op(e3, e1), e1)) & ~ (op(e1, e3) = op(e3, e1))) | (~ (e0 = op(op(e3, e0), e3)) & (e3 = op(op(e3, e0), e0)) & ~ (op(e0, e3) = op(e3, e0))) | (~ (e4 = op(op(e2, e4), e2)) & (e2 = op(op(e2, e4), e4)) & ~ (op(e2, e4) = op(e4, e2))) | (~ (e3 = op(op(e2, e3), e2)) & (e2 = op(op(e2, e3), e3)) & ~ (op(e2, e3) = op(e3, e2))) | (~ (e2 = op(op(e2, e2), e2)) & (e2 = op(op(e2, e2), e2)) & ~ (op(e2, e2) = op(e2, e2))) | (~ (e1 = op(op(e2, e1), e2)) & (e2 = op(op(e2, e1), e1)) & ~ (op(e1, e2) = op(e2, e1))) | (~ (e0 = op(op(e2, e0), e2)) & (e2 = op(op(e2, e0), e0)) & ~ (op(e0, e2) = op(e2, e0))) | (~ (e4 = op(op(e1, e4), e1)) & (e1 = op(op(e1, e4), e4)) & ~ (op(e1, e4) = op(e4, e1))) | (~ (e3 = op(op(e1, e3), e1)) & (e1 = op(op(e1, e3), e3)) & ~ (op(e1, e3) = op(e3, e1))) | (~ (e2 = op(op(e1, e2), e1)) & (e1 = op(op(e1, e2), e2)) & ~ (op(e1, e2) = op(e2, e1))) | (~ (e1 = op(op(e1, e1), e1)) & (e1 = op(op(e1, e1), e1)) & ~ (op(e1, e1) = op(e1, e1))) | (~ (e0 = op(op(e1, e0), e1)) & (e1 = op(op(e1, e0), e0)) & ~ (op(e0, e1) = op(e1, e0))) | (~ (e4 = op(op(e0, e4), e0)) & (e0 = op(op(e0, e4), e4)) & ~ (op(e0, e4) = op(e4, e0))) | (~ (e3 = op(op(e0, e3), e0)) & (e0 = op(op(e0, e3), e3)) & ~ (op(e0, e3) = op(e3, e0))) | (~ (e2 = op(op(e0, e2), e0)) & (e0 = op(op(e0, e2), e2)) & ~ (op(e0, e2) = op(e2, e0))) | (~ (e1 = op(op(e0, e1), e0)) & (e0 = op(op(e0, e1), e1)) & ~ (op(e0, e1) = op(e1, e0))) | (~ (e0 = op(op(e0, e0), e0)) & (e0 = op(op(e0, e0), e0)) & ~ (op(e0, e0) = op(e0, e0))) | (~ (e4 = op(e4, e4)) & (e4 = op(e4, e4)) & (e4 = op(e4, e4))) | (~ (e4 = op(e4, e3)) & (e4 = op(e3, e3)) & (e3 = op(e4, e4))) | (~ (e4 = op(e4, e2)) & (e4 = op(e2, e2)) & (e2 = op(e4, e4))) | (~ (e4 = op(e4, e1)) & (e4 = op(e1, e1)) & (e1 = op(e4, e4))) | (~ (e4 = op(e4, e0)) & (op(e0, e0) = e4) & (e0 = op(e4, e4))) | (~ (e3 = op(e3, e4)) & (e3 = op(e4, e4)) & (e4 = op(e3, e3))) | (~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | (~ (e3 = op(e3, e2)) & (e3 = op(e2, e2)) & (e2 = op(e3, e3))) | (~ (e3 = op(e3, e1)) & (e3 = op(e1, e1)) & (e1 = op(e3, e3))) | (~ (e3 = op(e3, e0)) & (op(e0, e0) = e3) & (e0 = op(e3, e3))) | (~ (e2 = op(e2, e4)) & (e2 = op(e4, e4)) & (e4 = op(e2, e2))) | (~ (e2 = op(e2, e3)) & (e2 = op(e3, e3)) & (e3 = op(e2, e2))) | (~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | (~ (e2 = op(e2, e1)) & (e2 = op(e1, e1)) & (e1 = op(e2, e2))) | (~ (e2 = op(e2, e0)) & (op(e0, e0) = e2) & (e0 = op(e2, e2))) | (~ (e1 = op(e1, e4)) & (e1 = op(e4, e4)) & (e4 = op(e1, e1))) | (~ (e1 = op(e1, e3)) & (e1 = op(e3, e3)) & (e3 = op(e1, e1))) | (~ (e1 = op(e1, e2)) & (e1 = op(e2, e2)) & (e2 = op(e1, e1))) | (~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | (~ (e1 = op(e1, e0)) & (op(e0, e0) = e1) & (e0 = op(e1, e1))) | (~ (e0 = op(e0, e4)) & (e0 = op(e4, e4)) & (op(e0, e0) = e4)) | (~ (e0 = op(e0, e3)) & (e0 = op(e3, e3)) & (op(e0, e0) = e3)) | (~ (e0 = op(e0, e2)) & (e0 = op(e2, e2)) & (op(e0, e0) = e2)) | (~ (e0 = op(e0, e1)) & (e0 = op(e1, e1)) & (op(e0, e0) = e1)) | (~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | (((~ (e4 = unit) & (e4 = op(e4, e4))) | ~ (e4 = op(e4, e4))) & ((~ (e4 = unit) & (e3 = op(e4, e4))) | ~ (e4 = op(e4, e3))) & ((~ (e4 = unit) & (e2 = op(e4, e4))) | ~ (e4 = op(e4, e2))) & ((~ (e4 = unit) & (e1 = op(e4, e4))) | ~ (e4 = op(e4, e1))) & ((~ (e4 = unit) & (e0 = op(e4, e4))) | ~ (e4 = op(e4, e0))) & ((~ (e4 = unit) & (e4 = op(e3, e4))) | ~ (e4 = op(e3, e4))) & ((~ (e4 = unit) & (e3 = op(e3, e4))) | ~ (e4 = op(e3, e3))) & ((~ (e4 = unit) & (e2 = op(e3, e4))) | ~ (e4 = op(e3, e2))) & ((~ (e4 = unit) & (e1 = op(e3, e4))) | ~ (e4 = op(e3, e1))) & ((~ (e4 = unit) & (e0 = op(e3, e4))) | ~ (e4 = op(e3, e0))) & ((~ (e4 = unit) & (e4 = op(e2, e4))) | ~ (e4 = op(e2, e4))) & ((~ (e4 = unit) & (e3 = op(e2, e4))) | ~ (e4 = op(e2, e3))) & ((~ (e4 = unit) & (e2 = op(e2, e4))) | ~ (e4 = op(e2, e2))) & ((~ (e4 = unit) & (e1 = op(e2, e4))) | ~ (e4 = op(e2, e1))) & ((~ (e4 = unit) & (e0 = op(e2, e4))) | ~ (e4 = op(e2, e0))) & ((~ (e4 = unit) & (e4 = op(e1, e4))) | ~ (e4 = op(e1, e4))) & ((~ (e4 = unit) & (e3 = op(e1, e4))) | ~ (e4 = op(e1, e3))) & ((~ (e4 = unit) & (e2 = op(e1, e4))) | ~ (e4 = op(e1, e2))) & ((~ (e4 = unit) & (e1 = op(e1, e4))) | ~ (e4 = op(e1, e1))) & ((~ (e4 = unit) & (e0 = op(e1, e4))) | ~ (e4 = op(e1, e0))) & ((~ (e4 = unit) & (e4 = op(e0, e4))) | ~ (e4 = op(e0, e4))) & ((~ (e4 = unit) & (e3 = op(e0, e4))) | ~ (e4 = op(e0, e3))) & ((~ (e4 = unit) & (e2 = op(e0, e4))) | ~ (e4 = op(e0, e2))) & ((~ (e4 = unit) & (e1 = op(e0, e4))) | ~ (e4 = op(e0, e1))) & ((~ (e4 = unit) & (e0 = op(e0, e4))) | ~ (op(e0, e0) = e4))) | (((~ (e3 = unit) & (e4 = op(e4, e3))) | ~ (e3 = op(e4, e4))) & ((~ (e3 = unit) & (e3 = op(e4, e3))) | ~ (e3 = op(e4, e3))) & ((~ (e3 = unit) & (e2 = op(e4, e3))) | ~ (e3 = op(e4, e2))) & ((~ (e3 = unit) & (e1 = op(e4, e3))) | ~ (e3 = op(e4, e1))) & ((~ (e3 = unit) & (e0 = op(e4, e3))) | ~ (e3 = op(e4, e0))) & ((~ (e3 = unit) & (e4 = op(e3, e3))) | ~ (e3 = op(e3, e4))) & ((~ (e3 = unit) & (e3 = op(e3, e3))) | ~ (e3 = op(e3, e3))) & ((~ (e3 = unit) & (e2 = op(e3, e3))) | ~ (e3 = op(e3, e2))) & ((~ (e3 = unit) & (e1 = op(e3, e3))) | ~ (e3 = op(e3, e1))) & ((~ (e3 = unit) & (e0 = op(e3, e3))) | ~ (e3 = op(e3, e0))) & ((~ (e3 = unit) & (e4 = op(e2, e3))) | ~ (e3 = op(e2, e4))) & ((~ (e3 = unit) & (e3 = op(e2, e3))) | ~ (e3 = op(e2, e3))) & ((~ (e3 = unit) & (e2 = op(e2, e3))) | ~ (e3 = op(e2, e2))) & ((~ (e3 = unit) & (e1 = op(e2, e3))) | ~ (e3 = op(e2, e1))) & ((~ (e3 = unit) & (e0 = op(e2, e3))) | ~ (e3 = op(e2, e0))) & ((~ (e3 = unit) & (e4 = op(e1, e3))) | ~ (e3 = op(e1, e4))) & ((~ (e3 = unit) & (e3 = op(e1, e3))) | ~ (e3 = op(e1, e3))) & ((~ (e3 = unit) & (e2 = op(e1, e3))) | ~ (e3 = op(e1, e2))) & ((~ (e3 = unit) & (e1 = op(e1, e3))) | ~ (e3 = op(e1, e1))) & ((~ (e3 = unit) & (e0 = op(e1, e3))) | ~ (e3 = op(e1, e0))) & ((~ (e3 = unit) & (e4 = op(e0, e3))) | ~ (e3 = op(e0, e4))) & ((~ (e3 = unit) & (e3 = op(e0, e3))) | ~ (e3 = op(e0, e3))) & ((~ (e3 = unit) & (e2 = op(e0, e3))) | ~ (e3 = op(e0, e2))) & ((~ (e3 = unit) & (e1 = op(e0, e3))) | ~ (e3 = op(e0, e1))) & ((~ (e3 = unit) & (e0 = op(e0, e3))) | ~ (op(e0, e0) = e3))) | (((~ (e2 = unit) & (e4 = op(e4, e2))) | ~ (e2 = op(e4, e4))) & ((~ (e2 = unit) & (e3 = op(e4, e2))) | ~ (e2 = op(e4, e3))) & ((~ (e2 = unit) & (e2 = op(e4, e2))) | ~ (e2 = op(e4, e2))) & ((~ (e2 = unit) & (e1 = op(e4, e2))) | ~ (e2 = op(e4, e1))) & ((~ (e2 = unit) & (e0 = op(e4, e2))) | ~ (e2 = op(e4, e0))) & ((~ (e2 = unit) & (e4 = op(e3, e2))) | ~ (e2 = op(e3, e4))) & ((~ (e2 = unit) & (e3 = op(e3, e2))) | ~ (e2 = op(e3, e3))) & ((~ (e2 = unit) & (e2 = op(e3, e2))) | ~ (e2 = op(e3, e2))) & ((~ (e2 = unit) & (e1 = op(e3, e2))) | ~ (e2 = op(e3, e1))) & ((~ (e2 = unit) & (e0 = op(e3, e2))) | ~ (e2 = op(e3, e0))) & ((~ (e2 = unit) & (e4 = op(e2, e2))) | ~ (e2 = op(e2, e4))) & ((~ (e2 = unit) & (e3 = op(e2, e2))) | ~ (e2 = op(e2, e3))) & ((~ (e2 = unit) & (e2 = op(e2, e2))) | ~ (e2 = op(e2, e2))) & ((~ (e2 = unit) & (e1 = op(e2, e2))) | ~ (e2 = op(e2, e1))) & ((~ (e2 = unit) & (e0 = op(e2, e2))) | ~ (e2 = op(e2, e0))) & ((~ (e2 = unit) & (e4 = op(e1, e2))) | ~ (e2 = op(e1, e4))) & ((~ (e2 = unit) & (e3 = op(e1, e2))) | ~ (e2 = op(e1, e3))) & ((~ (e2 = unit) & (e2 = op(e1, e2))) | ~ (e2 = op(e1, e2))) & ((~ (e2 = unit) & (e1 = op(e1, e2))) | ~ (e2 = op(e1, e1))) & ((~ (e2 = unit) & (e0 = op(e1, e2))) | ~ (e2 = op(e1, e0))) & ((~ (e2 = unit) & (e4 = op(e0, e2))) | ~ (e2 = op(e0, e4))) & ((~ (e2 = unit) & (e3 = op(e0, e2))) | ~ (e2 = op(e0, e3))) & ((~ (e2 = unit) & (e2 = op(e0, e2))) | ~ (e2 = op(e0, e2))) & ((~ (e2 = unit) & (e1 = op(e0, e2))) | ~ (e2 = op(e0, e1))) & ((~ (e2 = unit) & (e0 = op(e0, e2))) | ~ (op(e0, e0) = e2))) | (((~ (e1 = unit) & (e4 = op(e4, e1))) | ~ (e1 = op(e4, e4))) & ((~ (e1 = unit) & (e3 = op(e4, e1))) | ~ (e1 = op(e4, e3))) & ((~ (e1 = unit) & (e2 = op(e4, e1))) | ~ (e1 = op(e4, e2))) & ((~ (e1 = unit) & (e1 = op(e4, e1))) | ~ (e1 = op(e4, e1))) & ((~ (e1 = unit) & (e0 = op(e4, e1))) | ~ (e1 = op(e4, e0))) & ((~ (e1 = unit) & (e4 = op(e3, e1))) | ~ (e1 = op(e3, e4))) & ((~ (e1 = unit) & (e3 = op(e3, e1))) | ~ (e1 = op(e3, e3))) & ((~ (e1 = unit) & (e2 = op(e3, e1))) | ~ (e1 = op(e3, e2))) & ((~ (e1 = unit) & (e1 = op(e3, e1))) | ~ (e1 = op(e3, e1))) & ((~ (e1 = unit) & (e0 = op(e3, e1))) | ~ (e1 = op(e3, e0))) & ((~ (e1 = unit) & (e4 = op(e2, e1))) | ~ (e1 = op(e2, e4))) & ((~ (e1 = unit) & (e3 = op(e2, e1))) | ~ (e1 = op(e2, e3))) & ((~ (e1 = unit) & (e2 = op(e2, e1))) | ~ (e1 = op(e2, e2))) & ((~ (e1 = unit) & (e1 = op(e2, e1))) | ~ (e1 = op(e2, e1))) & ((~ (e1 = unit) & (e0 = op(e2, e1))) | ~ (e1 = op(e2, e0))) & ((~ (e1 = unit) & (e4 = op(e1, e1))) | ~ (e1 = op(e1, e4))) & ((~ (e1 = unit) & (e3 = op(e1, e1))) | ~ (e1 = op(e1, e3))) & ((~ (e1 = unit) & (e2 = op(e1, e1))) | ~ (e1 = op(e1, e2))) & ((~ (e1 = unit) & (e1 = op(e1, e1))) | ~ (e1 = op(e1, e1))) & ((~ (e1 = unit) & (e0 = op(e1, e1))) | ~ (e1 = op(e1, e0))) & ((~ (e1 = unit) & (e4 = op(e0, e1))) | ~ (e1 = op(e0, e4))) & ((~ (e1 = unit) & (e3 = op(e0, e1))) | ~ (e1 = op(e0, e3))) & ((~ (e1 = unit) & (e2 = op(e0, e1))) | ~ (e1 = op(e0, e2))) & ((~ (e1 = unit) & (e1 = op(e0, e1))) | ~ (e1 = op(e0, e1))) & ((~ (e1 = unit) & (e0 = op(e0, e1))) | ~ (op(e0, e0) = e1))) | (((~ (e0 = unit) & (e4 = op(e4, e0))) | ~ (e0 = op(e4, e4))) & ((~ (e0 = unit) & (e3 = op(e4, e0))) | ~ (e0 = op(e4, e3))) & ((~ (e0 = unit) & (e2 = op(e4, e0))) | ~ (e0 = op(e4, e2))) & ((~ (e0 = unit) & (e1 = op(e4, e0))) | ~ (e0 = op(e4, e1))) & ((~ (e0 = unit) & (e0 = op(e4, e0))) | ~ (e0 = op(e4, e0))) & ((~ (e0 = unit) & (e4 = op(e3, e0))) | ~ (e0 = op(e3, e4))) & ((~ (e0 = unit) & (e3 = op(e3, e0))) | ~ (e0 = op(e3, e3))) & ((~ (e0 = unit) & (e2 = op(e3, e0))) | ~ (e0 = op(e3, e2))) & ((~ (e0 = unit) & (e1 = op(e3, e0))) | ~ (e0 = op(e3, e1))) & ((~ (e0 = unit) & (e0 = op(e3, e0))) | ~ (e0 = op(e3, e0))) & ((~ (e0 = unit) & (e4 = op(e2, e0))) | ~ (e0 = op(e2, e4))) & ((~ (e0 = unit) & (e3 = op(e2, e0))) | ~ (e0 = op(e2, e3))) & ((~ (e0 = unit) & (e2 = op(e2, e0))) | ~ (e0 = op(e2, e2))) & ((~ (e0 = unit) & (e1 = op(e2, e0))) | ~ (e0 = op(e2, e1))) & ((~ (e0 = unit) & (e0 = op(e2, e0))) | ~ (e0 = op(e2, e0))) & ((~ (e0 = unit) & (e4 = op(e1, e0))) | ~ (e0 = op(e1, e4))) & ((~ (e0 = unit) & (e3 = op(e1, e0))) | ~ (e0 = op(e1, e3))) & ((~ (e0 = unit) & (e2 = op(e1, e0))) | ~ (e0 = op(e1, e2))) & ((~ (e0 = unit) & (e1 = op(e1, e0))) | ~ (e0 = op(e1, e1))) & ((~ (e0 = unit) & (e0 = op(e1, e0))) | ~ (e0 = op(e1, e0))) & ((~ (e0 = unit) & (op(e0, e0) = e4)) | ~ (e0 = op(e0, e4))) & ((~ (e0 = unit) & (op(e0, e0) = e3)) | ~ (e0 = op(e0, e3))) & ((~ (e0 = unit) & (op(e0, e0) = e2)) | ~ (e0 = op(e0, e2))) & ((~ (e0 = unit) & (op(e0, e0) = e1)) | ~ (e0 = op(e0, e1))) & ((~ (e0 = unit) & (e0 = op(e0, e0))) | ~ (e0 = op(e0, e0)))) | ((((e4 = op(e4, op(e4, e4))) & ~ (e4 = op(e4, op(e4, e4)))) | ((e3 = op(e4, op(e3, e4))) & ~ (e4 = op(e3, op(e3, e4)))) | ((e2 = op(e4, op(e2, e4))) & ~ (e4 = op(e2, op(e2, e4)))) | ((e1 = op(e4, op(e1, e4))) & ~ (e4 = op(e1, op(e1, e4)))) | ((e0 = op(e4, op(e0, e4))) & ~ (e4 = op(e0, op(e0, e4))))) & (((e4 = op(e3, op(e4, e3))) & ~ (e3 = op(e4, op(e4, e3)))) | ((e3 = op(e3, op(e3, e3))) & ~ (e3 = op(e3, op(e3, e3)))) | ((e2 = op(e3, op(e2, e3))) & ~ (e3 = op(e2, op(e2, e3)))) | ((e1 = op(e3, op(e1, e3))) & ~ (e3 = op(e1, op(e1, e3)))) | ((e0 = op(e3, op(e0, e3))) & ~ (e3 = op(e0, op(e0, e3))))) & (((e4 = op(e2, op(e4, e2))) & ~ (e2 = op(e4, op(e4, e2)))) | ((e3 = op(e2, op(e3, e2))) & ~ (e2 = op(e3, op(e3, e2)))) | ((e2 = op(e2, op(e2, e2))) & ~ (e2 = op(e2, op(e2, e2)))) | ((e1 = op(e2, op(e1, e2))) & ~ (e2 = op(e1, op(e1, e2)))) | ((e0 = op(e2, op(e0, e2))) & ~ (e2 = op(e0, op(e0, e2))))) & (((e4 = op(e1, op(e4, e1))) & ~ (e1 = op(e4, op(e4, e1)))) | ((e3 = op(e1, op(e3, e1))) & ~ (e1 = op(e3, op(e3, e1)))) | ((e2 = op(e1, op(e2, e1))) & ~ (e1 = op(e2, op(e2, e1)))) | ((e1 = op(e1, op(e1, e1))) & ~ (e1 = op(e1, op(e1, e1)))) | ((e0 = op(e1, op(e0, e1))) & ~ (e1 = op(e0, op(e0, e1))))) & (((e4 = op(e0, op(e4, e0))) & ~ (e0 = op(e4, op(e4, e0)))) | ((e3 = op(e0, op(e3, e0))) & ~ (e0 = op(e3, op(e3, e0)))) | ((e2 = op(e0, op(e2, e0))) & ~ (e0 = op(e2, op(e2, e0)))) | ((e1 = op(e0, op(e1, e0))) & ~ (e0 = op(e1, op(e1, e0)))) | ((e0 = op(e0, op(e0, e0))) & ~ (e0 = op(e0, op(e0, e0)))))) | (((e4 = op(e4, e4)) | ~ (e4 = op(e4, e4))) & ((e4 = op(e4, e3)) | ~ (e3 = op(e4, e4))) & ((e4 = op(e4, e2)) | ~ (e2 = op(e4, e4))) & ((e4 = op(e4, e1)) | ~ (e1 = op(e4, e4))) & ((e4 = op(e4, e0)) | ~ (e0 = op(e4, e4))) & ((e3 = op(e3, e4)) | ~ (e4 = op(e3, e3))) & ((e3 = op(e3, e3)) | ~ (e3 = op(e3, e3))) & ((e3 = op(e3, e2)) | ~ (e2 = op(e3, e3))) & ((e3 = op(e3, e1)) | ~ (e1 = op(e3, e3))) & ((e3 = op(e3, e0)) | ~ (e0 = op(e3, e3))) & ((e2 = op(e2, e4)) | ~ (e4 = op(e2, e2))) & ((e2 = op(e2, e3)) | ~ (e3 = op(e2, e2))) & ((e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & ((e2 = op(e2, e1)) | ~ (e1 = op(e2, e2))) & ((e2 = op(e2, e0)) | ~ (e0 = op(e2, e2))) & ((e1 = op(e1, e4)) | ~ (e4 = op(e1, e1))) & ((e1 = op(e1, e3)) | ~ (e3 = op(e1, e1))) & ((e1 = op(e1, e2)) | ~ (e2 = op(e1, e1))) & ((e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & ((e1 = op(e1, e0)) | ~ (e0 = op(e1, e1))) & ((e0 = op(e0, e4)) | ~ (op(e0, e0) = e4)) & ((e0 = op(e0, e3)) | ~ (op(e0, e0) = e3)) & ((e0 = op(e0, e2)) | ~ (op(e0, e0) = e2)) & ((e0 = op(e0, e1)) | ~ (op(e0, e0) = e1)) & ((e0 = op(e0, e0)) | ~ (e0 = op(e0, e0)))) | (((e4 = op(e4, e4)) | (e4 = op(e3, e3)) | (e4 = op(e2, e2)) | (e4 = op(e1, e1)) | (op(e0, e0) = e4)) & ((e3 = op(e4, e4)) | (e3 = op(e3, e3)) | (e3 = op(e2, e2)) | (e3 = op(e1, e1)) | (op(e0, e0) = e3)) & ((e2 = op(e4, e4)) | (e2 = op(e3, e3)) | (e2 = op(e2, e2)) | (e2 = op(e1, e1)) | (op(e0, e0) = e2)) & ((e1 = op(e4, e4)) | (e1 = op(e3, e3)) | (e1 = op(e2, e2)) | (e1 = op(e1, e1)) | (op(e0, e0) = e1)) & ((e0 = op(e4, e4)) | (e0 = op(e3, e3)) | (e0 = op(e2, e2)) | (e0 = op(e1, e1)) | (e0 = op(e0, e0))))), inference(negated_conjecture, [], [f7])).
fof(f7, plain, ~ ((~ (e4 = op(op(e4, e4), e4)) & (e4 = op(op(e4, e4), e4)) & ~ (op(e4, e4) = op(e4, e4))) | (~ (e3 = op(op(e4, e3), e4)) & (e4 = op(op(e4, e3), e3)) & ~ (op(e3, e4) = op(e4, e3))) | (~ (e2 = op(op(e4, e2), e4)) & (e4 = op(op(e4, e2), e2)) & ~ (op(e2, e4) = op(e4, e2))) | (~ (e1 = op(op(e4, e1), e4)) & (e4 = op(op(e4, e1), e1)) & ~ (op(e1, e4) = op(e4, e1))) | (~ (e0 = op(op(e4, e0), e4)) & (e4 = op(op(e4, e0), e0)) & ~ (op(e0, e4) = op(e4, e0))) | (~ (e4 = op(op(e3, e4), e3)) & (e3 = op(op(e3, e4), e4)) & ~ (op(e3, e4) = op(e4, e3))) | (~ (e3 = op(op(e3, e3), e3)) & (e3 = op(op(e3, e3), e3)) & ~ (op(e3, e3) = op(e3, e3))) | (~ (e2 = op(op(e3, e2), e3)) & (e3 = op(op(e3, e2), e2)) & ~ (op(e2, e3) = op(e3, e2))) | (~ (e1 = op(op(e3, e1), e3)) & (e3 = op(op(e3, e1), e1)) & ~ (op(e1, e3) = op(e3, e1))) | (~ (e0 = op(op(e3, e0), e3)) & (e3 = op(op(e3, e0), e0)) & ~ (op(e0, e3) = op(e3, e0))) | (~ (e4 = op(op(e2, e4), e2)) & (e2 = op(op(e2, e4), e4)) & ~ (op(e2, e4) = op(e4, e2))) | (~ (e3 = op(op(e2, e3), e2)) & (e2 = op(op(e2, e3), e3)) & ~ (op(e2, e3) = op(e3, e2))) | (~ (e2 = op(op(e2, e2), e2)) & (e2 = op(op(e2, e2), e2)) & ~ (op(e2, e2) = op(e2, e2))) | (~ (e1 = op(op(e2, e1), e2)) & (e2 = op(op(e2, e1), e1)) & ~ (op(e1, e2) = op(e2, e1))) | (~ (e0 = op(op(e2, e0), e2)) & (e2 = op(op(e2, e0), e0)) & ~ (op(e0, e2) = op(e2, e0))) | (~ (e4 = op(op(e1, e4), e1)) & (e1 = op(op(e1, e4), e4)) & ~ (op(e1, e4) = op(e4, e1))) | (~ (e3 = op(op(e1, e3), e1)) & (e1 = op(op(e1, e3), e3)) & ~ (op(e1, e3) = op(e3, e1))) | (~ (e2 = op(op(e1, e2), e1)) & (e1 = op(op(e1, e2), e2)) & ~ (op(e1, e2) = op(e2, e1))) | (~ (e1 = op(op(e1, e1), e1)) & (e1 = op(op(e1, e1), e1)) & ~ (op(e1, e1) = op(e1, e1))) | (~ (e0 = op(op(e1, e0), e1)) & (e1 = op(op(e1, e0), e0)) & ~ (op(e0, e1) = op(e1, e0))) | (~ (e4 = op(op(e0, e4), e0)) & (e0 = op(op(e0, e4), e4)) & ~ (op(e0, e4) = op(e4, e0))) | (~ (e3 = op(op(e0, e3), e0)) & (e0 = op(op(e0, e3), e3)) & ~ (op(e0, e3) = op(e3, e0))) | (~ (e2 = op(op(e0, e2), e0)) & (e0 = op(op(e0, e2), e2)) & ~ (op(e0, e2) = op(e2, e0))) | (~ (e1 = op(op(e0, e1), e0)) & (e0 = op(op(e0, e1), e1)) & ~ (op(e0, e1) = op(e1, e0))) | (~ (e0 = op(op(e0, e0), e0)) & (e0 = op(op(e0, e0), e0)) & ~ (op(e0, e0) = op(e0, e0))) | (~ (e4 = op(e4, e4)) & (e4 = op(e4, e4)) & (e4 = op(e4, e4))) | (~ (e4 = op(e4, e3)) & (e4 = op(e3, e3)) & (e3 = op(e4, e4))) | (~ (e4 = op(e4, e2)) & (e4 = op(e2, e2)) & (e2 = op(e4, e4))) | (~ (e4 = op(e4, e1)) & (e4 = op(e1, e1)) & (e1 = op(e4, e4))) | (~ (e4 = op(e4, e0)) & (op(e0, e0) = e4) & (e0 = op(e4, e4))) | (~ (e3 = op(e3, e4)) & (e3 = op(e4, e4)) & (e4 = op(e3, e3))) | (~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | (~ (e3 = op(e3, e2)) & (e3 = op(e2, e2)) & (e2 = op(e3, e3))) | (~ (e3 = op(e3, e1)) & (e3 = op(e1, e1)) & (e1 = op(e3, e3))) | (~ (e3 = op(e3, e0)) & (op(e0, e0) = e3) & (e0 = op(e3, e3))) | (~ (e2 = op(e2, e4)) & (e2 = op(e4, e4)) & (e4 = op(e2, e2))) | (~ (e2 = op(e2, e3)) & (e2 = op(e3, e3)) & (e3 = op(e2, e2))) | (~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | (~ (e2 = op(e2, e1)) & (e2 = op(e1, e1)) & (e1 = op(e2, e2))) | (~ (e2 = op(e2, e0)) & (op(e0, e0) = e2) & (e0 = op(e2, e2))) | (~ (e1 = op(e1, e4)) & (e1 = op(e4, e4)) & (e4 = op(e1, e1))) | (~ (e1 = op(e1, e3)) & (e1 = op(e3, e3)) & (e3 = op(e1, e1))) | (~ (e1 = op(e1, e2)) & (e1 = op(e2, e2)) & (e2 = op(e1, e1))) | (~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | (~ (e1 = op(e1, e0)) & (op(e0, e0) = e1) & (e0 = op(e1, e1))) | (~ (e0 = op(e0, e4)) & (e0 = op(e4, e4)) & (op(e0, e0) = e4)) | (~ (e0 = op(e0, e3)) & (e0 = op(e3, e3)) & (op(e0, e0) = e3)) | (~ (e0 = op(e0, e2)) & (e0 = op(e2, e2)) & (op(e0, e0) = e2)) | (~ (e0 = op(e0, e1)) & (e0 = op(e1, e1)) & (op(e0, e0) = e1)) | (~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | (((~ (e4 = unit) & (e4 = op(e4, e4))) | ~ (e4 = op(e4, e4))) & ((~ (e4 = unit) & (e3 = op(e4, e4))) | ~ (e4 = op(e4, e3))) & ((~ (e4 = unit) & (e2 = op(e4, e4))) | ~ (e4 = op(e4, e2))) & ((~ (e4 = unit) & (e1 = op(e4, e4))) | ~ (e4 = op(e4, e1))) & ((~ (e4 = unit) & (e0 = op(e4, e4))) | ~ (e4 = op(e4, e0))) & ((~ (e4 = unit) & (e4 = op(e3, e4))) | ~ (e4 = op(e3, e4))) & ((~ (e4 = unit) & (e3 = op(e3, e4))) | ~ (e4 = op(e3, e3))) & ((~ (e4 = unit) & (e2 = op(e3, e4))) | ~ (e4 = op(e3, e2))) & ((~ (e4 = unit) & (e1 = op(e3, e4))) | ~ (e4 = op(e3, e1))) & ((~ (e4 = unit) & (e0 = op(e3, e4))) | ~ (e4 = op(e3, e0))) & ((~ (e4 = unit) & (e4 = op(e2, e4))) | ~ (e4 = op(e2, e4))) & ((~ (e4 = unit) & (e3 = op(e2, e4))) | ~ (e4 = op(e2, e3))) & ((~ (e4 = unit) & (e2 = op(e2, e4))) | ~ (e4 = op(e2, e2))) & ((~ (e4 = unit) & (e1 = op(e2, e4))) | ~ (e4 = op(e2, e1))) & ((~ (e4 = unit) & (e0 = op(e2, e4))) | ~ (e4 = op(e2, e0))) & ((~ (e4 = unit) & (e4 = op(e1, e4))) | ~ (e4 = op(e1, e4))) & ((~ (e4 = unit) & (e3 = op(e1, e4))) | ~ (e4 = op(e1, e3))) & ((~ (e4 = unit) & (e2 = op(e1, e4))) | ~ (e4 = op(e1, e2))) & ((~ (e4 = unit) & (e1 = op(e1, e4))) | ~ (e4 = op(e1, e1))) & ((~ (e4 = unit) & (e0 = op(e1, e4))) | ~ (e4 = op(e1, e0))) & ((~ (e4 = unit) & (e4 = op(e0, e4))) | ~ (e4 = op(e0, e4))) & ((~ (e4 = unit) & (e3 = op(e0, e4))) | ~ (e4 = op(e0, e3))) & ((~ (e4 = unit) & (e2 = op(e0, e4))) | ~ (e4 = op(e0, e2))) & ((~ (e4 = unit) & (e1 = op(e0, e4))) | ~ (e4 = op(e0, e1))) & ((~ (e4 = unit) & (e0 = op(e0, e4))) | ~ (op(e0, e0) = e4))) | (((~ (e3 = unit) & (e4 = op(e4, e3))) | ~ (e3 = op(e4, e4))) & ((~ (e3 = unit) & (e3 = op(e4, e3))) | ~ (e3 = op(e4, e3))) & ((~ (e3 = unit) & (e2 = op(e4, e3))) | ~ (e3 = op(e4, e2))) & ((~ (e3 = unit) & (e1 = op(e4, e3))) | ~ (e3 = op(e4, e1))) & ((~ (e3 = unit) & (e0 = op(e4, e3))) | ~ (e3 = op(e4, e0))) & ((~ (e3 = unit) & (e4 = op(e3, e3))) | ~ (e3 = op(e3, e4))) & ((~ (e3 = unit) & (e3 = op(e3, e3))) | ~ (e3 = op(e3, e3))) & ((~ (e3 = unit) & (e2 = op(e3, e3))) | ~ (e3 = op(e3, e2))) & ((~ (e3 = unit) & (e1 = op(e3, e3))) | ~ (e3 = op(e3, e1))) & ((~ (e3 = unit) & (e0 = op(e3, e3))) | ~ (e3 = op(e3, e0))) & ((~ (e3 = unit) & (e4 = op(e2, e3))) | ~ (e3 = op(e2, e4))) & ((~ (e3 = unit) & (e3 = op(e2, e3))) | ~ (e3 = op(e2, e3))) & ((~ (e3 = unit) & (e2 = op(e2, e3))) | ~ (e3 = op(e2, e2))) & ((~ (e3 = unit) & (e1 = op(e2, e3))) | ~ (e3 = op(e2, e1))) & ((~ (e3 = unit) & (e0 = op(e2, e3))) | ~ (e3 = op(e2, e0))) & ((~ (e3 = unit) & (e4 = op(e1, e3))) | ~ (e3 = op(e1, e4))) & ((~ (e3 = unit) & (e3 = op(e1, e3))) | ~ (e3 = op(e1, e3))) & ((~ (e3 = unit) & (e2 = op(e1, e3))) | ~ (e3 = op(e1, e2))) & ((~ (e3 = unit) & (e1 = op(e1, e3))) | ~ (e3 = op(e1, e1))) & ((~ (e3 = unit) & (e0 = op(e1, e3))) | ~ (e3 = op(e1, e0))) & ((~ (e3 = unit) & (e4 = op(e0, e3))) | ~ (e3 = op(e0, e4))) & ((~ (e3 = unit) & (e3 = op(e0, e3))) | ~ (e3 = op(e0, e3))) & ((~ (e3 = unit) & (e2 = op(e0, e3))) | ~ (e3 = op(e0, e2))) & ((~ (e3 = unit) & (e1 = op(e0, e3))) | ~ (e3 = op(e0, e1))) & ((~ (e3 = unit) & (e0 = op(e0, e3))) | ~ (op(e0, e0) = e3))) | (((~ (e2 = unit) & (e4 = op(e4, e2))) | ~ (e2 = op(e4, e4))) & ((~ (e2 = unit) & (e3 = op(e4, e2))) | ~ (e2 = op(e4, e3))) & ((~ (e2 = unit) & (e2 = op(e4, e2))) | ~ (e2 = op(e4, e2))) & ((~ (e2 = unit) & (e1 = op(e4, e2))) | ~ (e2 = op(e4, e1))) & ((~ (e2 = unit) & (e0 = op(e4, e2))) | ~ (e2 = op(e4, e0))) & ((~ (e2 = unit) & (e4 = op(e3, e2))) | ~ (e2 = op(e3, e4))) & ((~ (e2 = unit) & (e3 = op(e3, e2))) | ~ (e2 = op(e3, e3))) & ((~ (e2 = unit) & (e2 = op(e3, e2))) | ~ (e2 = op(e3, e2))) & ((~ (e2 = unit) & (e1 = op(e3, e2))) | ~ (e2 = op(e3, e1))) & ((~ (e2 = unit) & (e0 = op(e3, e2))) | ~ (e2 = op(e3, e0))) & ((~ (e2 = unit) & (e4 = op(e2, e2))) | ~ (e2 = op(e2, e4))) & ((~ (e2 = unit) & (e3 = op(e2, e2))) | ~ (e2 = op(e2, e3))) & ((~ (e2 = unit) & (e2 = op(e2, e2))) | ~ (e2 = op(e2, e2))) & ((~ (e2 = unit) & (e1 = op(e2, e2))) | ~ (e2 = op(e2, e1))) & ((~ (e2 = unit) & (e0 = op(e2, e2))) | ~ (e2 = op(e2, e0))) & ((~ (e2 = unit) & (e4 = op(e1, e2))) | ~ (e2 = op(e1, e4))) & ((~ (e2 = unit) & (e3 = op(e1, e2))) | ~ (e2 = op(e1, e3))) & ((~ (e2 = unit) & (e2 = op(e1, e2))) | ~ (e2 = op(e1, e2))) & ((~ (e2 = unit) & (e1 = op(e1, e2))) | ~ (e2 = op(e1, e1))) & ((~ (e2 = unit) & (e0 = op(e1, e2))) | ~ (e2 = op(e1, e0))) & ((~ (e2 = unit) & (e4 = op(e0, e2))) | ~ (e2 = op(e0, e4))) & ((~ (e2 = unit) & (e3 = op(e0, e2))) | ~ (e2 = op(e0, e3))) & ((~ (e2 = unit) & (e2 = op(e0, e2))) | ~ (e2 = op(e0, e2))) & ((~ (e2 = unit) & (e1 = op(e0, e2))) | ~ (e2 = op(e0, e1))) & ((~ (e2 = unit) & (e0 = op(e0, e2))) | ~ (op(e0, e0) = e2))) | (((~ (e1 = unit) & (e4 = op(e4, e1))) | ~ (e1 = op(e4, e4))) & ((~ (e1 = unit) & (e3 = op(e4, e1))) | ~ (e1 = op(e4, e3))) & ((~ (e1 = unit) & (e2 = op(e4, e1))) | ~ (e1 = op(e4, e2))) & ((~ (e1 = unit) & (e1 = op(e4, e1))) | ~ (e1 = op(e4, e1))) & ((~ (e1 = unit) & (e0 = op(e4, e1))) | ~ (e1 = op(e4, e0))) & ((~ (e1 = unit) & (e4 = op(e3, e1))) | ~ (e1 = op(e3, e4))) & ((~ (e1 = unit) & (e3 = op(e3, e1))) | ~ (e1 = op(e3, e3))) & ((~ (e1 = unit) & (e2 = op(e3, e1))) | ~ (e1 = op(e3, e2))) & ((~ (e1 = unit) & (e1 = op(e3, e1))) | ~ (e1 = op(e3, e1))) & ((~ (e1 = unit) & (e0 = op(e3, e1))) | ~ (e1 = op(e3, e0))) & ((~ (e1 = unit) & (e4 = op(e2, e1))) | ~ (e1 = op(e2, e4))) & ((~ (e1 = unit) & (e3 = op(e2, e1))) | ~ (e1 = op(e2, e3))) & ((~ (e1 = unit) & (e2 = op(e2, e1))) | ~ (e1 = op(e2, e2))) & ((~ (e1 = unit) & (e1 = op(e2, e1))) | ~ (e1 = op(e2, e1))) & ((~ (e1 = unit) & (e0 = op(e2, e1))) | ~ (e1 = op(e2, e0))) & ((~ (e1 = unit) & (e4 = op(e1, e1))) | ~ (e1 = op(e1, e4))) & ((~ (e1 = unit) & (e3 = op(e1, e1))) | ~ (e1 = op(e1, e3))) & ((~ (e1 = unit) & (e2 = op(e1, e1))) | ~ (e1 = op(e1, e2))) & ((~ (e1 = unit) & (e1 = op(e1, e1))) | ~ (e1 = op(e1, e1))) & ((~ (e1 = unit) & (e0 = op(e1, e1))) | ~ (e1 = op(e1, e0))) & ((~ (e1 = unit) & (e4 = op(e0, e1))) | ~ (e1 = op(e0, e4))) & ((~ (e1 = unit) & (e3 = op(e0, e1))) | ~ (e1 = op(e0, e3))) & ((~ (e1 = unit) & (e2 = op(e0, e1))) | ~ (e1 = op(e0, e2))) & ((~ (e1 = unit) & (e1 = op(e0, e1))) | ~ (e1 = op(e0, e1))) & ((~ (e1 = unit) & (e0 = op(e0, e1))) | ~ (op(e0, e0) = e1))) | (((~ (e0 = unit) & (e4 = op(e4, e0))) | ~ (e0 = op(e4, e4))) & ((~ (e0 = unit) & (e3 = op(e4, e0))) | ~ (e0 = op(e4, e3))) & ((~ (e0 = unit) & (e2 = op(e4, e0))) | ~ (e0 = op(e4, e2))) & ((~ (e0 = unit) & (e1 = op(e4, e0))) | ~ (e0 = op(e4, e1))) & ((~ (e0 = unit) & (e0 = op(e4, e0))) | ~ (e0 = op(e4, e0))) & ((~ (e0 = unit) & (e4 = op(e3, e0))) | ~ (e0 = op(e3, e4))) & ((~ (e0 = unit) & (e3 = op(e3, e0))) | ~ (e0 = op(e3, e3))) & ((~ (e0 = unit) & (e2 = op(e3, e0))) | ~ (e0 = op(e3, e2))) & ((~ (e0 = unit) & (e1 = op(e3, e0))) | ~ (e0 = op(e3, e1))) & ((~ (e0 = unit) & (e0 = op(e3, e0))) | ~ (e0 = op(e3, e0))) & ((~ (e0 = unit) & (e4 = op(e2, e0))) | ~ (e0 = op(e2, e4))) & ((~ (e0 = unit) & (e3 = op(e2, e0))) | ~ (e0 = op(e2, e3))) & ((~ (e0 = unit) & (e2 = op(e2, e0))) | ~ (e0 = op(e2, e2))) & ((~ (e0 = unit) & (e1 = op(e2, e0))) | ~ (e0 = op(e2, e1))) & ((~ (e0 = unit) & (e0 = op(e2, e0))) | ~ (e0 = op(e2, e0))) & ((~ (e0 = unit) & (e4 = op(e1, e0))) | ~ (e0 = op(e1, e4))) & ((~ (e0 = unit) & (e3 = op(e1, e0))) | ~ (e0 = op(e1, e3))) & ((~ (e0 = unit) & (e2 = op(e1, e0))) | ~ (e0 = op(e1, e2))) & ((~ (e0 = unit) & (e1 = op(e1, e0))) | ~ (e0 = op(e1, e1))) & ((~ (e0 = unit) & (e0 = op(e1, e0))) | ~ (e0 = op(e1, e0))) & ((~ (e0 = unit) & (op(e0, e0) = e4)) | ~ (e0 = op(e0, e4))) & ((~ (e0 = unit) & (op(e0, e0) = e3)) | ~ (e0 = op(e0, e3))) & ((~ (e0 = unit) & (op(e0, e0) = e2)) | ~ (e0 = op(e0, e2))) & ((~ (e0 = unit) & (op(e0, e0) = e1)) | ~ (e0 = op(e0, e1))) & ((~ (e0 = unit) & (e0 = op(e0, e0))) | ~ (e0 = op(e0, e0)))) | ((((e4 = op(e4, op(e4, e4))) & ~ (e4 = op(e4, op(e4, e4)))) | ((e3 = op(e4, op(e3, e4))) & ~ (e4 = op(e3, op(e3, e4)))) | ((e2 = op(e4, op(e2, e4))) & ~ (e4 = op(e2, op(e2, e4)))) | ((e1 = op(e4, op(e1, e4))) & ~ (e4 = op(e1, op(e1, e4)))) | ((e0 = op(e4, op(e0, e4))) & ~ (e4 = op(e0, op(e0, e4))))) & (((e4 = op(e3, op(e4, e3))) & ~ (e3 = op(e4, op(e4, e3)))) | ((e3 = op(e3, op(e3, e3))) & ~ (e3 = op(e3, op(e3, e3)))) | ((e2 = op(e3, op(e2, e3))) & ~ (e3 = op(e2, op(e2, e3)))) | ((e1 = op(e3, op(e1, e3))) & ~ (e3 = op(e1, op(e1, e3)))) | ((e0 = op(e3, op(e0, e3))) & ~ (e3 = op(e0, op(e0, e3))))) & (((e4 = op(e2, op(e4, e2))) & ~ (e2 = op(e4, op(e4, e2)))) | ((e3 = op(e2, op(e3, e2))) & ~ (e2 = op(e3, op(e3, e2)))) | ((e2 = op(e2, op(e2, e2))) & ~ (e2 = op(e2, op(e2, e2)))) | ((e1 = op(e2, op(e1, e2))) & ~ (e2 = op(e1, op(e1, e2)))) | ((e0 = op(e2, op(e0, e2))) & ~ (e2 = op(e0, op(e0, e2))))) & (((e4 = op(e1, op(e4, e1))) & ~ (e1 = op(e4, op(e4, e1)))) | ((e3 = op(e1, op(e3, e1))) & ~ (e1 = op(e3, op(e3, e1)))) | ((e2 = op(e1, op(e2, e1))) & ~ (e1 = op(e2, op(e2, e1)))) | ((e1 = op(e1, op(e1, e1))) & ~ (e1 = op(e1, op(e1, e1)))) | ((e0 = op(e1, op(e0, e1))) & ~ (e1 = op(e0, op(e0, e1))))) & (((e4 = op(e0, op(e4, e0))) & ~ (e0 = op(e4, op(e4, e0)))) | ((e3 = op(e0, op(e3, e0))) & ~ (e0 = op(e3, op(e3, e0)))) | ((e2 = op(e0, op(e2, e0))) & ~ (e0 = op(e2, op(e2, e0)))) | ((e1 = op(e0, op(e1, e0))) & ~ (e0 = op(e1, op(e1, e0)))) | ((e0 = op(e0, op(e0, e0))) & ~ (e0 = op(e0, op(e0, e0)))))) | (((e4 = op(e4, e4)) | ~ (e4 = op(e4, e4))) & ((e4 = op(e4, e3)) | ~ (e3 = op(e4, e4))) & ((e4 = op(e4, e2)) | ~ (e2 = op(e4, e4))) & ((e4 = op(e4, e1)) | ~ (e1 = op(e4, e4))) & ((e4 = op(e4, e0)) | ~ (e0 = op(e4, e4))) & ((e3 = op(e3, e4)) | ~ (e4 = op(e3, e3))) & ((e3 = op(e3, e3)) | ~ (e3 = op(e3, e3))) & ((e3 = op(e3, e2)) | ~ (e2 = op(e3, e3))) & ((e3 = op(e3, e1)) | ~ (e1 = op(e3, e3))) & ((e3 = op(e3, e0)) | ~ (e0 = op(e3, e3))) & ((e2 = op(e2, e4)) | ~ (e4 = op(e2, e2))) & ((e2 = op(e2, e3)) | ~ (e3 = op(e2, e2))) & ((e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & ((e2 = op(e2, e1)) | ~ (e1 = op(e2, e2))) & ((e2 = op(e2, e0)) | ~ (e0 = op(e2, e2))) & ((e1 = op(e1, e4)) | ~ (e4 = op(e1, e1))) & ((e1 = op(e1, e3)) | ~ (e3 = op(e1, e1))) & ((e1 = op(e1, e2)) | ~ (e2 = op(e1, e1))) & ((e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & ((e1 = op(e1, e0)) | ~ (e0 = op(e1, e1))) & ((e0 = op(e0, e4)) | ~ (op(e0, e0) = e4)) & ((e0 = op(e0, e3)) | ~ (op(e0, e0) = e3)) & ((e0 = op(e0, e2)) | ~ (op(e0, e0) = e2)) & ((e0 = op(e0, e1)) | ~ (op(e0, e0) = e1)) & ((e0 = op(e0, e0)) | ~ (e0 = op(e0, e0)))) | (((e4 = op(e4, e4)) | (e4 = op(e3, e3)) | (e4 = op(e2, e2)) | (e4 = op(e1, e1)) | (op(e0, e0) = e4)) & ((e3 = op(e4, e4)) | (e3 = op(e3, e3)) | (e3 = op(e2, e2)) | (e3 = op(e1, e1)) | (op(e0, e0) = e3)) & ((e2 = op(e4, e4)) | (e2 = op(e3, e3)) | (e2 = op(e2, e2)) | (e2 = op(e1, e1)) | (op(e0, e0) = e2)) & ((e1 = op(e4, e4)) | (e1 = op(e3, e3)) | (e1 = op(e2, e2)) | (e1 = op(e1, e1)) | (op(e0, e0) = e1)) & ((e0 = op(e4, e4)) | (e0 = op(e3, e3)) | (e0 = op(e2, e2)) | (e0 = op(e1, e1)) | (e0 = op(e0, e0))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG047+1.p', co1)).
fof(f3223, plain, (spl140_297 | spl140_296 | spl140_295 | spl140_294 | ~ spl140_65), inference(avatar_split_clause, [], [f794, f1506, f2841, f2850, f2859, f2868])).
fof(f794, plain, (~ (e4 = op(e2, e2)) | sP3 | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f150])).
fof(f3222, plain, (spl140_297 | spl140_296 | spl140_295 | spl140_294 | ~ spl140_35), inference(avatar_split_clause, [], [f795, f1380, f2841, f2850, f2859, f2868])).
fof(f795, plain, (~ (e4 = op(e3, e3)) | sP3 | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f150])).
fof(f3220, plain, (spl140_293 | spl140_292 | spl140_291 | spl140_290 | spl140_289 | spl140_288 | spl140_287 | spl140_286 | spl140_285 | spl140_284 | spl140_283 | spl140_282 | spl140_281 | spl140_280 | spl140_279 | spl140_278 | spl140_277 | spl140_276 | spl140_275 | spl140_274 | spl140_273 | spl140_272 | spl140_3 | spl140_4 | spl140_5), inference(avatar_split_clause, [], [f797, f1254, f1250, f1246, f2709, f2715, f2721, f2727, f2733, f2739, f2745, f2751, f2757, f2763, f2769, f2775, f2781, f2787, f2793, f2799, f2805, f2811, f2817, f2823, f2829, f2835])).
fof(f2835, plain, (spl140_293 <=> sP4), introduced(avatar_definition, [new_symbols(naming, [spl140_293])])).
fof(f2829, plain, (spl140_292 <=> sP5), introduced(avatar_definition, [new_symbols(naming, [spl140_292])])).
fof(f2823, plain, (spl140_291 <=> sP6), introduced(avatar_definition, [new_symbols(naming, [spl140_291])])).
fof(f2817, plain, (spl140_290 <=> sP7), introduced(avatar_definition, [new_symbols(naming, [spl140_290])])).
fof(f2811, plain, (spl140_289 <=> sP8), introduced(avatar_definition, [new_symbols(naming, [spl140_289])])).
fof(f2805, plain, (spl140_288 <=> sP9), introduced(avatar_definition, [new_symbols(naming, [spl140_288])])).
fof(f2799, plain, (spl140_287 <=> sP10), introduced(avatar_definition, [new_symbols(naming, [spl140_287])])).
fof(f2793, plain, (spl140_286 <=> sP11), introduced(avatar_definition, [new_symbols(naming, [spl140_286])])).
fof(f2787, plain, (spl140_285 <=> sP12), introduced(avatar_definition, [new_symbols(naming, [spl140_285])])).
fof(f2781, plain, (spl140_284 <=> sP13), introduced(avatar_definition, [new_symbols(naming, [spl140_284])])).
fof(f2775, plain, (spl140_283 <=> sP14), introduced(avatar_definition, [new_symbols(naming, [spl140_283])])).
fof(f2769, plain, (spl140_282 <=> sP15), introduced(avatar_definition, [new_symbols(naming, [spl140_282])])).
fof(f2763, plain, (spl140_281 <=> sP16), introduced(avatar_definition, [new_symbols(naming, [spl140_281])])).
fof(f2757, plain, (spl140_280 <=> sP17), introduced(avatar_definition, [new_symbols(naming, [spl140_280])])).
fof(f2751, plain, (spl140_279 <=> sP18), introduced(avatar_definition, [new_symbols(naming, [spl140_279])])).
fof(f2745, plain, (spl140_278 <=> sP19), introduced(avatar_definition, [new_symbols(naming, [spl140_278])])).
fof(f2739, plain, (spl140_277 <=> sP20), introduced(avatar_definition, [new_symbols(naming, [spl140_277])])).
fof(f2733, plain, (spl140_276 <=> sP21), introduced(avatar_definition, [new_symbols(naming, [spl140_276])])).
fof(f2727, plain, (spl140_275 <=> sP22), introduced(avatar_definition, [new_symbols(naming, [spl140_275])])).
fof(f2721, plain, (spl140_274 <=> sP23), introduced(avatar_definition, [new_symbols(naming, [spl140_274])])).
fof(f2715, plain, (spl140_273 <=> sP24), introduced(avatar_definition, [new_symbols(naming, [spl140_273])])).
fof(f2709, plain, (spl140_272 <=> sP25), introduced(avatar_definition, [new_symbols(naming, [spl140_272])])).
fof(f797, plain, ((e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4), inference(cnf_transformation, [], [f150])).
fof(f3203, plain, (spl140_263 | spl140_254 | spl140_245 | spl140_236 | spl140_352 | ~ spl140_353), inference(avatar_split_clause, [], [f806, f3200, f3196, f2549, f2589, f2629, f2669])).
fof(f2669, plain, (spl140_263 <=> sP26), introduced(avatar_definition, [new_symbols(naming, [spl140_263])])).
fof(f2629, plain, (spl140_254 <=> sP27), introduced(avatar_definition, [new_symbols(naming, [spl140_254])])).
fof(f2589, plain, (spl140_245 <=> sP28), introduced(avatar_definition, [new_symbols(naming, [spl140_245])])).
fof(f2549, plain, (spl140_236 <=> sP29), introduced(avatar_definition, [new_symbols(naming, [spl140_236])])).
fof(f806, plain, (~ (e1 = op(e4, op(e1, e4))) | (e4 = op(e1, op(e1, e4))) | sP29 | sP28 | sP27 | sP26), inference(cnf_transformation, [], [f150])).
fof(f3194, plain, (spl140_263 | spl140_254 | spl140_245 | spl140_236 | spl140_350 | ~ spl140_351), inference(avatar_split_clause, [], [f807, f3191, f3187, f2549, f2589, f2629, f2669])).
fof(f807, plain, (~ (e2 = op(e4, op(e2, e4))) | (e4 = op(e2, op(e2, e4))) | sP29 | sP28 | sP27 | sP26), inference(cnf_transformation, [], [f150])).
fof(f3185, plain, (spl140_263 | spl140_254 | spl140_245 | spl140_236 | spl140_348 | ~ spl140_349), inference(avatar_split_clause, [], [f808, f3182, f3178, f2549, f2589, f2629, f2669])).
fof(f808, plain, (~ (e3 = op(e4, op(e3, e4))) | (e4 = op(e3, op(e3, e4))) | sP29 | sP28 | sP27 | sP26), inference(cnf_transformation, [], [f150])).
fof(f3174, plain, (spl140_213 | spl140_212 | spl140_211 | spl140_210 | spl140_209 | spl140_208 | spl140_207 | spl140_206 | spl140_205 | spl140_204 | spl140_203 | spl140_202 | spl140_201 | spl140_200 | spl140_199 | spl140_198 | spl140_197 | spl140_196 | spl140_195 | spl140_194 | spl140_193 | spl140_192 | ~ spl140_18 | spl140_7 | spl140_2), inference(avatar_split_clause, [], [f3173, f1242, f1263, f1309, f2286, f2293, f2300, f2307, f2314, f2321, f2328, f2335, f2342, f2349, f2356, f2363, f2370, f2377, f2384, f2391, f2398, f2405, f2412, f2419, f2426, f2433])).
fof(f2433, plain, (spl140_213 <=> sP52), introduced(avatar_definition, [new_symbols(naming, [spl140_213])])).
fof(f2426, plain, (spl140_212 <=> sP53), introduced(avatar_definition, [new_symbols(naming, [spl140_212])])).
fof(f2419, plain, (spl140_211 <=> sP54), introduced(avatar_definition, [new_symbols(naming, [spl140_211])])).
fof(f2412, plain, (spl140_210 <=> sP55), introduced(avatar_definition, [new_symbols(naming, [spl140_210])])).
fof(f2405, plain, (spl140_209 <=> sP56), introduced(avatar_definition, [new_symbols(naming, [spl140_209])])).
fof(f2398, plain, (spl140_208 <=> sP57), introduced(avatar_definition, [new_symbols(naming, [spl140_208])])).
fof(f2391, plain, (spl140_207 <=> sP58), introduced(avatar_definition, [new_symbols(naming, [spl140_207])])).
fof(f2384, plain, (spl140_206 <=> sP59), introduced(avatar_definition, [new_symbols(naming, [spl140_206])])).
fof(f2377, plain, (spl140_205 <=> sP60), introduced(avatar_definition, [new_symbols(naming, [spl140_205])])).
fof(f2370, plain, (spl140_204 <=> sP61), introduced(avatar_definition, [new_symbols(naming, [spl140_204])])).
fof(f2363, plain, (spl140_203 <=> sP62), introduced(avatar_definition, [new_symbols(naming, [spl140_203])])).
fof(f2356, plain, (spl140_202 <=> sP63), introduced(avatar_definition, [new_symbols(naming, [spl140_202])])).
fof(f2349, plain, (spl140_201 <=> sP64), introduced(avatar_definition, [new_symbols(naming, [spl140_201])])).
fof(f2342, plain, (spl140_200 <=> sP65), introduced(avatar_definition, [new_symbols(naming, [spl140_200])])).
fof(f2335, plain, (spl140_199 <=> sP66), introduced(avatar_definition, [new_symbols(naming, [spl140_199])])).
fof(f2328, plain, (spl140_198 <=> sP67), introduced(avatar_definition, [new_symbols(naming, [spl140_198])])).
fof(f2321, plain, (spl140_197 <=> sP68), introduced(avatar_definition, [new_symbols(naming, [spl140_197])])).
fof(f2314, plain, (spl140_196 <=> sP69), introduced(avatar_definition, [new_symbols(naming, [spl140_196])])).
fof(f2307, plain, (spl140_195 <=> sP70), introduced(avatar_definition, [new_symbols(naming, [spl140_195])])).
fof(f2300, plain, (spl140_194 <=> sP71), introduced(avatar_definition, [new_symbols(naming, [spl140_194])])).
fof(f2293, plain, (spl140_193 <=> sP72), introduced(avatar_definition, [new_symbols(naming, [spl140_193])])).
fof(f2286, plain, (spl140_192 <=> sP73), introduced(avatar_definition, [new_symbols(naming, [spl140_192])])).
fof(f3173, plain, ((e1 = op(e4, e4)) | (e1 = op(e4, e3)) | ~ (e2 = op(e4, e1)) | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52), inference(subsumption_resolution, [], [f819, f1033])).
fof(f819, plain, ((e1 = op(e4, e4)) | (e1 = op(e4, e3)) | (e1 = unit) | ~ (e2 = op(e4, e1)) | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52), inference(cnf_transformation, [], [f150])).
fof(f3172, plain, (spl140_213 | spl140_212 | spl140_211 | spl140_210 | spl140_209 | spl140_208 | spl140_207 | spl140_206 | spl140_205 | spl140_204 | spl140_203 | spl140_202 | spl140_201 | spl140_200 | spl140_199 | spl140_198 | spl140_197 | spl140_196 | spl140_195 | spl140_194 | spl140_193 | spl140_192 | spl140_12 | ~ spl140_19 | spl140_2), inference(avatar_split_clause, [], [f3171, f1242, f1313, f1284, f2286, f2293, f2300, f2307, f2314, f2321, f2328, f2335, f2342, f2349, f2356, f2363, f2370, f2377, f2384, f2391, f2398, f2405, f2412, f2419, f2426, f2433])).
fof(f3171, plain, ((e1 = op(e4, e4)) | ~ (e3 = op(e4, e1)) | (e1 = op(e4, e2)) | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52), inference(subsumption_resolution, [], [f820, f1033])).
fof(f820, plain, ((e1 = op(e4, e4)) | (e1 = unit) | ~ (e3 = op(e4, e1)) | (e1 = op(e4, e2)) | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52), inference(cnf_transformation, [], [f150])).
fof(f3160, plain, (spl140_191 | spl140_190 | spl140_189 | spl140_188 | spl140_187 | spl140_186 | spl140_185 | spl140_184 | spl140_183 | spl140_182 | spl140_181 | spl140_180 | spl140_179 | spl140_178 | spl140_177 | spl140_176 | spl140_175 | spl140_174 | spl140_173 | spl140_172 | spl140_171 | spl140_170 | spl140_13 | spl140_8 | spl140_3), inference(avatar_split_clause, [], [f826, f1246, f1267, f1288, f2132, f2139, f2146, f2153, f2160, f2167, f2174, f2181, f2188, f2195, f2202, f2209, f2216, f2223, f2230, f2237, f2244, f2251, f2258, f2265, f2272, f2279])).
fof(f2279, plain, (spl140_191 <=> sP74), introduced(avatar_definition, [new_symbols(naming, [spl140_191])])).
fof(f2272, plain, (spl140_190 <=> sP75), introduced(avatar_definition, [new_symbols(naming, [spl140_190])])).
fof(f2265, plain, (spl140_189 <=> sP76), introduced(avatar_definition, [new_symbols(naming, [spl140_189])])).
fof(f2258, plain, (spl140_188 <=> sP77), introduced(avatar_definition, [new_symbols(naming, [spl140_188])])).
fof(f2251, plain, (spl140_187 <=> sP78), introduced(avatar_definition, [new_symbols(naming, [spl140_187])])).
fof(f2244, plain, (spl140_186 <=> sP79), introduced(avatar_definition, [new_symbols(naming, [spl140_186])])).
fof(f2237, plain, (spl140_185 <=> sP80), introduced(avatar_definition, [new_symbols(naming, [spl140_185])])).
fof(f2230, plain, (spl140_184 <=> sP81), introduced(avatar_definition, [new_symbols(naming, [spl140_184])])).
fof(f2223, plain, (spl140_183 <=> sP82), introduced(avatar_definition, [new_symbols(naming, [spl140_183])])).
fof(f2216, plain, (spl140_182 <=> sP83), introduced(avatar_definition, [new_symbols(naming, [spl140_182])])).
fof(f2209, plain, (spl140_181 <=> sP84), introduced(avatar_definition, [new_symbols(naming, [spl140_181])])).
fof(f2202, plain, (spl140_180 <=> sP85), introduced(avatar_definition, [new_symbols(naming, [spl140_180])])).
fof(f2195, plain, (spl140_179 <=> sP86), introduced(avatar_definition, [new_symbols(naming, [spl140_179])])).
fof(f2188, plain, (spl140_178 <=> sP87), introduced(avatar_definition, [new_symbols(naming, [spl140_178])])).
fof(f2181, plain, (spl140_177 <=> sP88), introduced(avatar_definition, [new_symbols(naming, [spl140_177])])).
fof(f2174, plain, (spl140_176 <=> sP89), introduced(avatar_definition, [new_symbols(naming, [spl140_176])])).
fof(f2167, plain, (spl140_175 <=> sP90), introduced(avatar_definition, [new_symbols(naming, [spl140_175])])).
fof(f2160, plain, (spl140_174 <=> sP91), introduced(avatar_definition, [new_symbols(naming, [spl140_174])])).
fof(f2153, plain, (spl140_173 <=> sP92), introduced(avatar_definition, [new_symbols(naming, [spl140_173])])).
fof(f2146, plain, (spl140_172 <=> sP93), introduced(avatar_definition, [new_symbols(naming, [spl140_172])])).
fof(f2139, plain, (spl140_171 <=> sP94), introduced(avatar_definition, [new_symbols(naming, [spl140_171])])).
fof(f2132, plain, (spl140_170 <=> sP95), introduced(avatar_definition, [new_symbols(naming, [spl140_170])])).
fof(f826, plain, ((e2 = op(e4, e4)) | (e2 = op(e4, e3)) | (e2 = op(e4, e2)) | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74), inference(cnf_transformation, [], [f150])).
fof(f3157, plain, (spl140_191 | spl140_190 | spl140_189 | spl140_188 | spl140_187 | spl140_186 | spl140_185 | spl140_184 | spl140_183 | spl140_182 | spl140_181 | spl140_180 | spl140_179 | spl140_178 | spl140_177 | spl140_176 | spl140_175 | spl140_174 | spl140_173 | spl140_172 | spl140_171 | spl140_170 | spl140_13 | ~ spl140_14 | spl140_3), inference(avatar_split_clause, [], [f3156, f1246, f1292, f1288, f2132, f2139, f2146, f2153, f2160, f2167, f2174, f2181, f2188, f2195, f2202, f2209, f2216, f2223, f2230, f2237, f2244, f2251, f2258, f2265, f2272, f2279])).
fof(f3156, plain, ((e2 = op(e4, e4)) | ~ (e3 = op(e4, e2)) | (e2 = op(e4, e2)) | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74), inference(subsumption_resolution, [], [f828, f1032])).
fof(f828, plain, ((e2 = op(e4, e4)) | (e2 = unit) | ~ (e3 = op(e4, e2)) | (e2 = op(e4, e2)) | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74), inference(cnf_transformation, [], [f150])).
fof(f3145, plain, (spl140_169 | spl140_168 | spl140_167 | spl140_166 | spl140_165 | spl140_164 | spl140_163 | spl140_162 | spl140_161 | spl140_160 | spl140_159 | spl140_158 | spl140_157 | spl140_156 | spl140_155 | spl140_154 | spl140_153 | spl140_152 | spl140_151 | spl140_150 | spl140_149 | spl140_148 | spl140_14 | spl140_9 | spl140_4), inference(avatar_split_clause, [], [f834, f1250, f1271, f1292, f1978, f1985, f1992, f1999, f2006, f2013, f2020, f2027, f2034, f2041, f2048, f2055, f2062, f2069, f2076, f2083, f2090, f2097, f2104, f2111, f2118, f2125])).
fof(f2125, plain, (spl140_169 <=> sP96), introduced(avatar_definition, [new_symbols(naming, [spl140_169])])).
fof(f2118, plain, (spl140_168 <=> sP97), introduced(avatar_definition, [new_symbols(naming, [spl140_168])])).
fof(f2111, plain, (spl140_167 <=> sP98), introduced(avatar_definition, [new_symbols(naming, [spl140_167])])).
fof(f2104, plain, (spl140_166 <=> sP99), introduced(avatar_definition, [new_symbols(naming, [spl140_166])])).
fof(f2097, plain, (spl140_165 <=> sP100), introduced(avatar_definition, [new_symbols(naming, [spl140_165])])).
fof(f2090, plain, (spl140_164 <=> sP101), introduced(avatar_definition, [new_symbols(naming, [spl140_164])])).
fof(f2083, plain, (spl140_163 <=> sP102), introduced(avatar_definition, [new_symbols(naming, [spl140_163])])).
fof(f2076, plain, (spl140_162 <=> sP103), introduced(avatar_definition, [new_symbols(naming, [spl140_162])])).
fof(f2069, plain, (spl140_161 <=> sP104), introduced(avatar_definition, [new_symbols(naming, [spl140_161])])).
fof(f2062, plain, (spl140_160 <=> sP105), introduced(avatar_definition, [new_symbols(naming, [spl140_160])])).
fof(f2055, plain, (spl140_159 <=> sP106), introduced(avatar_definition, [new_symbols(naming, [spl140_159])])).
fof(f2048, plain, (spl140_158 <=> sP107), introduced(avatar_definition, [new_symbols(naming, [spl140_158])])).
fof(f2041, plain, (spl140_157 <=> sP108), introduced(avatar_definition, [new_symbols(naming, [spl140_157])])).
fof(f2034, plain, (spl140_156 <=> sP109), introduced(avatar_definition, [new_symbols(naming, [spl140_156])])).
fof(f2027, plain, (spl140_155 <=> sP110), introduced(avatar_definition, [new_symbols(naming, [spl140_155])])).
fof(f2020, plain, (spl140_154 <=> sP111), introduced(avatar_definition, [new_symbols(naming, [spl140_154])])).
fof(f2013, plain, (spl140_153 <=> sP112), introduced(avatar_definition, [new_symbols(naming, [spl140_153])])).
fof(f2006, plain, (spl140_152 <=> sP113), introduced(avatar_definition, [new_symbols(naming, [spl140_152])])).
fof(f1999, plain, (spl140_151 <=> sP114), introduced(avatar_definition, [new_symbols(naming, [spl140_151])])).
fof(f1992, plain, (spl140_150 <=> sP115), introduced(avatar_definition, [new_symbols(naming, [spl140_150])])).
fof(f1985, plain, (spl140_149 <=> sP116), introduced(avatar_definition, [new_symbols(naming, [spl140_149])])).
fof(f1978, plain, (spl140_148 <=> sP117), introduced(avatar_definition, [new_symbols(naming, [spl140_148])])).
fof(f834, plain, ((e3 = op(e4, e4)) | (e3 = op(e4, e3)) | (e3 = op(e4, e2)) | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96), inference(cnf_transformation, [], [f150])).
fof(f3144, plain, (spl140_169 | spl140_168 | spl140_167 | spl140_166 | spl140_165 | spl140_164 | spl140_163 | spl140_162 | spl140_161 | spl140_160 | spl140_159 | spl140_158 | spl140_157 | spl140_156 | spl140_155 | spl140_154 | spl140_153 | spl140_152 | spl140_151 | spl140_150 | spl140_149 | spl140_148 | ~ spl140_8 | spl140_9 | spl140_4), inference(avatar_split_clause, [], [f3143, f1250, f1271, f1267, f1978, f1985, f1992, f1999, f2006, f2013, f2020, f2027, f2034, f2041, f2048, f2055, f2062, f2069, f2076, f2083, f2090, f2097, f2104, f2111, f2118, f2125])).
fof(f3143, plain, ((e3 = op(e4, e4)) | (e3 = op(e4, e3)) | ~ (e2 = op(e4, e3)) | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96), inference(subsumption_resolution, [], [f835, f1031])).
fof(f835, plain, ((e3 = op(e4, e4)) | (e3 = op(e4, e3)) | (e3 = unit) | ~ (e2 = op(e4, e3)) | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96), inference(cnf_transformation, [], [f150])).
fof(f3130, plain, (spl140_147 | spl140_146 | spl140_145 | spl140_144 | spl140_143 | spl140_142 | spl140_141 | spl140_140 | spl140_139 | spl140_138 | spl140_137 | spl140_136 | spl140_135 | spl140_134 | spl140_133 | spl140_132 | spl140_131 | spl140_130 | spl140_129 | spl140_128 | spl140_127 | spl140_126 | spl140_15 | spl140_10 | spl140_5), inference(avatar_split_clause, [], [f842, f1254, f1275, f1296, f1824, f1831, f1838, f1845, f1852, f1859, f1866, f1873, f1880, f1887, f1894, f1901, f1908, f1915, f1922, f1929, f1936, f1943, f1950, f1957, f1964, f1971])).
fof(f1971, plain, (spl140_147 <=> sP118), introduced(avatar_definition, [new_symbols(naming, [spl140_147])])).
fof(f1964, plain, (spl140_146 <=> sP119), introduced(avatar_definition, [new_symbols(naming, [spl140_146])])).
fof(f1957, plain, (spl140_145 <=> sP120), introduced(avatar_definition, [new_symbols(naming, [spl140_145])])).
fof(f1950, plain, (spl140_144 <=> sP121), introduced(avatar_definition, [new_symbols(naming, [spl140_144])])).
fof(f1943, plain, (spl140_143 <=> sP122), introduced(avatar_definition, [new_symbols(naming, [spl140_143])])).
fof(f1936, plain, (spl140_142 <=> sP123), introduced(avatar_definition, [new_symbols(naming, [spl140_142])])).
fof(f1929, plain, (spl140_141 <=> sP124), introduced(avatar_definition, [new_symbols(naming, [spl140_141])])).
fof(f1922, plain, (spl140_140 <=> sP125), introduced(avatar_definition, [new_symbols(naming, [spl140_140])])).
fof(f1915, plain, (spl140_139 <=> sP126), introduced(avatar_definition, [new_symbols(naming, [spl140_139])])).
fof(f1908, plain, (spl140_138 <=> sP127), introduced(avatar_definition, [new_symbols(naming, [spl140_138])])).
fof(f1901, plain, (spl140_137 <=> sP128), introduced(avatar_definition, [new_symbols(naming, [spl140_137])])).
fof(f1894, plain, (spl140_136 <=> sP129), introduced(avatar_definition, [new_symbols(naming, [spl140_136])])).
fof(f1887, plain, (spl140_135 <=> sP130), introduced(avatar_definition, [new_symbols(naming, [spl140_135])])).
fof(f1880, plain, (spl140_134 <=> sP131), introduced(avatar_definition, [new_symbols(naming, [spl140_134])])).
fof(f1873, plain, (spl140_133 <=> sP132), introduced(avatar_definition, [new_symbols(naming, [spl140_133])])).
fof(f1866, plain, (spl140_132 <=> sP133), introduced(avatar_definition, [new_symbols(naming, [spl140_132])])).
fof(f1859, plain, (spl140_131 <=> sP134), introduced(avatar_definition, [new_symbols(naming, [spl140_131])])).
fof(f1852, plain, (spl140_130 <=> sP135), introduced(avatar_definition, [new_symbols(naming, [spl140_130])])).
fof(f1845, plain, (spl140_129 <=> sP136), introduced(avatar_definition, [new_symbols(naming, [spl140_129])])).
fof(f1838, plain, (spl140_128 <=> sP137), introduced(avatar_definition, [new_symbols(naming, [spl140_128])])).
fof(f1831, plain, (spl140_127 <=> sP138), introduced(avatar_definition, [new_symbols(naming, [spl140_127])])).
fof(f1824, plain, (spl140_126 <=> sP139), introduced(avatar_definition, [new_symbols(naming, [spl140_126])])).
fof(f842, plain, ((e4 = op(e4, e4)) | (e4 = op(e4, e3)) | (e4 = op(e4, e2)) | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118), inference(cnf_transformation, [], [f150])).
fof(f3115, plain, (~ spl140_122 | ~ spl140_91 | spl140_116), inference(avatar_split_clause, [], [f1200, f1721, f1616, f1746])).
fof(f1746, plain, (spl140_122 <=> (e1 = op(unit, unit))), introduced(avatar_definition, [new_symbols(naming, [spl140_122])])).
fof(f1200, plain, ((unit = op(unit, e1)) | ~ (op(e1, e1) = unit) | ~ (e1 = op(unit, unit))), inference(definition_unfolding, [], [f851, f377, f377, f377, f377, f377])).
fof(f851, plain, ((e0 = op(e0, e1)) | ~ (e0 = op(e1, e1)) | ~ (op(e0, e0) = e1)), inference(cnf_transformation, [], [f150])).
fof(f3110, plain, (~ spl140_93 | ~ spl140_62 | spl140_87), inference(avatar_split_clause, [], [f857, f1599, f1494, f1624])).
fof(f857, plain, ((e1 = op(e1, e2)) | ~ (e1 = op(e2, e2)) | ~ (e2 = op(e1, e1))), inference(cnf_transformation, [], [f150])).
fof(f3109, plain, (~ spl140_94 | ~ spl140_32 | spl140_82), inference(avatar_split_clause, [], [f858, f1578, f1368, f1628])).
fof(f858, plain, ((e1 = op(e1, e3)) | ~ (e1 = op(e3, e3)) | ~ (e3 = op(e1, e1))), inference(cnf_transformation, [], [f150])).
fof(f3108, plain, (~ spl140_95 | ~ spl140_2 | spl140_77), inference(avatar_split_clause, [], [f859, f1557, f1242, f1632])).
fof(f859, plain, ((e1 = op(e1, e4)) | ~ (e1 = op(e4, e4)) | ~ (e4 = op(e1, e1))), inference(cnf_transformation, [], [f150])).
fof(f3105, plain, (~ spl140_64 | ~ spl140_33 | spl140_58), inference(avatar_split_clause, [], [f863, f1477, f1372, f1502])).
fof(f863, plain, ((e2 = op(e2, e3)) | ~ (e2 = op(e3, e3)) | ~ (e3 = op(e2, e2))), inference(cnf_transformation, [], [f150])).
fof(f3104, plain, (~ spl140_65 | ~ spl140_3 | spl140_53), inference(avatar_split_clause, [], [f864, f1456, f1246, f1506])).
fof(f864, plain, ((e2 = op(e2, e4)) | ~ (e2 = op(e4, e4)) | ~ (e4 = op(e2, e2))), inference(cnf_transformation, [], [f150])).
fof(f3100, plain, (~ spl140_35 | ~ spl140_4 | spl140_29), inference(avatar_split_clause, [], [f869, f1355, f1250, f1380])).
fof(f869, plain, ((e3 = op(e3, e4)) | ~ (e3 = op(e4, e4)) | ~ (e4 = op(e3, e3))), inference(cnf_transformation, [], [f150])).
fof(f3037, plain, (spl140_315 | ~ spl140_333 | spl140_334), inference(avatar_split_clause, [], [f883, f3034, f3030, f2951])).
fof(f883, plain, ((e3 = op(op(e1, e3), e1)) | ~ (e1 = op(op(e1, e3), e3)) | (op(e1, e3) = op(e3, e1))), inference(cnf_transformation, [], [f150])).
fof(f3028, plain, (spl140_304 | ~ spl140_331 | spl140_332), inference(avatar_split_clause, [], [f884, f3025, f3021, f2903])).
fof(f884, plain, ((e4 = op(op(e1, e4), e1)) | ~ (e1 = op(op(e1, e4), e4)) | (op(e1, e4) = op(e4, e1))), inference(cnf_transformation, [], [f150])).
fof(f2993, plain, (spl140_312 | ~ spl140_323 | spl140_324), inference(avatar_split_clause, [], [f888, f2990, f2986, f2938])).
fof(f888, plain, ((e3 = op(op(e2, e3), e2)) | ~ (e2 = op(op(e2, e3), e3)) | (op(e2, e3) = op(e3, e2))), inference(cnf_transformation, [], [f150])).
fof(f2914, plain, (spl140_304 | ~ spl140_305 | spl140_306), inference(avatar_split_clause, [], [f896, f2911, f2907, f2903])).
fof(f896, plain, ((e1 = op(op(e4, e1), e4)) | ~ (e4 = op(op(e4, e1), e1)) | (op(e1, e4) = op(e4, e1))), inference(cnf_transformation, [], [f150])).
fof(f2901, plain, (spl140_301 | ~ spl140_302 | spl140_303), inference(avatar_split_clause, [], [f897, f2898, f2894, f2890])).
fof(f897, plain, ((e2 = op(op(e4, e2), e4)) | ~ (e4 = op(op(e4, e2), e2)) | (op(e2, e4) = op(e4, e2))), inference(cnf_transformation, [], [f150])).
fof(f2888, plain, (spl140_298 | ~ spl140_299 | spl140_300), inference(avatar_split_clause, [], [f898, f2885, f2881, f2877])).
fof(f898, plain, ((e3 = op(op(e4, e3), e4)) | ~ (e4 = op(op(e4, e3), e3)) | (op(e3, e4) = op(e4, e3))), inference(cnf_transformation, [], [f150])).
fof(f2875, plain, (~ spl140_297 | ~ spl140_121), inference(avatar_split_clause, [], [f1183, f1742, f2868])).
fof(f1183, plain, (~ (unit = op(unit, unit)) | ~ sP0), inference(definition_unfolding, [], [f787, f377, f377, f377])).
fof(f787, plain, (~ (e0 = op(e0, e0)) | ~ sP0), inference(cnf_transformation, [], [f290])).
fof(f290, plain, ((~ (e0 = op(e4, e4)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e0, e0))) | ~ sP0), inference(nnf_transformation, [], [f10])).
fof(f2866, plain, (~ spl140_296 | ~ spl140_122), inference(avatar_split_clause, [], [f1178, f1746, f2859])).
fof(f1178, plain, (~ (e1 = op(unit, unit)) | ~ sP1), inference(definition_unfolding, [], [f782, f377, f377])).
fof(f782, plain, (~ (op(e0, e0) = e1) | ~ sP1), inference(cnf_transformation, [], [f289])).
fof(f289, plain, ((~ (e1 = op(e4, e4)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e1, e1)) & ~ (op(e0, e0) = e1)) | ~ sP1), inference(nnf_transformation, [], [f11])).
fof(f2864, plain, (~ spl140_296 | ~ spl140_62), inference(avatar_split_clause, [], [f784, f1494, f2859])).
fof(f784, plain, (~ (e1 = op(e2, e2)) | ~ sP1), inference(cnf_transformation, [], [f289])).
fof(f2863, plain, (~ spl140_296 | ~ spl140_32), inference(avatar_split_clause, [], [f785, f1368, f2859])).
fof(f785, plain, (~ (e1 = op(e3, e3)) | ~ sP1), inference(cnf_transformation, [], [f289])).
fof(f2862, plain, (~ spl140_296 | ~ spl140_2), inference(avatar_split_clause, [], [f786, f1242, f2859])).
fof(f786, plain, (~ (e1 = op(e4, e4)) | ~ sP1), inference(cnf_transformation, [], [f289])).
fof(f2856, plain, (~ spl140_295 | ~ spl140_93), inference(avatar_split_clause, [], [f778, f1624, f2850])).
fof(f778, plain, (~ (e2 = op(e1, e1)) | ~ sP2), inference(cnf_transformation, [], [f288])).
fof(f288, plain, ((~ (e2 = op(e4, e4)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e1, e1)) & ~ (op(e0, e0) = e2)) | ~ sP2), inference(nnf_transformation, [], [f12])).
fof(f2854, plain, (~ spl140_295 | ~ spl140_33), inference(avatar_split_clause, [], [f780, f1372, f2850])).
fof(f780, plain, (~ (e2 = op(e3, e3)) | ~ sP2), inference(cnf_transformation, [], [f288])).
fof(f2853, plain, (~ spl140_295 | ~ spl140_3), inference(avatar_split_clause, [], [f781, f1246, f2850])).
fof(f781, plain, (~ (e2 = op(e4, e4)) | ~ sP2), inference(cnf_transformation, [], [f288])).
fof(f2847, plain, (~ spl140_294 | ~ spl140_94), inference(avatar_split_clause, [], [f773, f1628, f2841])).
fof(f773, plain, (~ (e3 = op(e1, e1)) | ~ sP3), inference(cnf_transformation, [], [f287])).
fof(f287, plain, ((~ (e3 = op(e4, e4)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e1, e1)) & ~ (op(e0, e0) = e3)) | ~ sP3), inference(nnf_transformation, [], [f13])).
fof(f2846, plain, (~ spl140_294 | ~ spl140_64), inference(avatar_split_clause, [], [f774, f1502, f2841])).
fof(f774, plain, (~ (e3 = op(e2, e2)) | ~ sP3), inference(cnf_transformation, [], [f287])).
fof(f2844, plain, (~ spl140_294 | ~ spl140_4), inference(avatar_split_clause, [], [f776, f1250, f2841])).
fof(f776, plain, (~ (e3 = op(e4, e4)) | ~ sP3), inference(cnf_transformation, [], [f287])).
fof(f2838, plain, (~ spl140_293 | ~ spl140_121), inference(avatar_split_clause, [], [f1174, f1742, f2835])).
fof(f1174, plain, (~ (unit = op(unit, unit)) | ~ sP4), inference(definition_unfolding, [], [f771, f377, f377, f377])).
fof(f771, plain, (~ (e0 = op(e0, e0)) | ~ sP4), inference(cnf_transformation, [], [f286])).
fof(f286, plain, ((~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP4), inference(nnf_transformation, [], [f14])).
fof(f2833, plain, (~ spl140_292 | spl140_122), inference(avatar_split_clause, [], [f1173, f1746, f2829])).
fof(f1173, plain, ((e1 = op(unit, unit)) | ~ sP5), inference(definition_unfolding, [], [f768, f377, f377])).
fof(f768, plain, ((op(e0, e0) = e1) | ~ sP5), inference(cnf_transformation, [], [f285])).
fof(f285, plain, ((~ (e0 = op(e0, e1)) & (op(e0, e0) = e1)) | ~ sP5), inference(nnf_transformation, [], [f15])).
fof(f2827, plain, (~ spl140_291 | spl140_123), inference(avatar_split_clause, [], [f1171, f1750, f2823])).
fof(f1171, plain, ((e2 = op(unit, unit)) | ~ sP6), inference(definition_unfolding, [], [f766, f377, f377])).
fof(f766, plain, ((op(e0, e0) = e2) | ~ sP6), inference(cnf_transformation, [], [f284])).
fof(f284, plain, ((~ (e0 = op(e0, e2)) & (op(e0, e0) = e2)) | ~ sP6), inference(nnf_transformation, [], [f16])).
fof(f2821, plain, (~ spl140_290 | spl140_124), inference(avatar_split_clause, [], [f1169, f1754, f2817])).
fof(f1169, plain, ((e3 = op(unit, unit)) | ~ sP7), inference(definition_unfolding, [], [f764, f377, f377])).
fof(f764, plain, ((op(e0, e0) = e3) | ~ sP7), inference(cnf_transformation, [], [f283])).
fof(f283, plain, ((~ (e0 = op(e0, e3)) & (op(e0, e0) = e3)) | ~ sP7), inference(nnf_transformation, [], [f17])).
fof(f2815, plain, (~ spl140_289 | spl140_125), inference(avatar_split_clause, [], [f1167, f1758, f2811])).
fof(f1167, plain, ((e4 = op(unit, unit)) | ~ sP8), inference(definition_unfolding, [], [f762, f377, f377])).
fof(f762, plain, ((op(e0, e0) = e4) | ~ sP8), inference(cnf_transformation, [], [f282])).
fof(f282, plain, ((~ (e0 = op(e0, e4)) & (op(e0, e0) = e4)) | ~ sP8), inference(nnf_transformation, [], [f18])).
fof(f2808, plain, (~ spl140_288 | ~ spl140_97), inference(avatar_split_clause, [], [f1164, f1641, f2805])).
fof(f1164, plain, (~ (e1 = op(e1, unit)) | ~ sP9), inference(definition_unfolding, [], [f761, f377])).
fof(f761, plain, (~ (e1 = op(e1, e0)) | ~ sP9), inference(cnf_transformation, [], [f281])).
fof(f281, plain, ((~ (e1 = op(e1, e0)) & (e0 = op(e1, e1))) | ~ sP9), inference(nnf_transformation, [], [f19])).
fof(f2803, plain, (~ spl140_287 | spl140_92), inference(avatar_split_clause, [], [f758, f1620, f2799])).
fof(f758, plain, ((e1 = op(e1, e1)) | ~ sP10), inference(cnf_transformation, [], [f280])).
fof(f280, plain, ((~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP10), inference(nnf_transformation, [], [f20])).
fof(f2797, plain, (~ spl140_286 | spl140_93), inference(avatar_split_clause, [], [f756, f1624, f2793])).
fof(f756, plain, ((e2 = op(e1, e1)) | ~ sP11), inference(cnf_transformation, [], [f279])).
fof(f279, plain, ((~ (e1 = op(e1, e2)) & (e2 = op(e1, e1))) | ~ sP11), inference(nnf_transformation, [], [f21])).
fof(f2791, plain, (~ spl140_285 | spl140_94), inference(avatar_split_clause, [], [f754, f1628, f2787])).
fof(f754, plain, ((e3 = op(e1, e1)) | ~ sP12), inference(cnf_transformation, [], [f278])).
fof(f278, plain, ((~ (e1 = op(e1, e3)) & (e3 = op(e1, e1))) | ~ sP12), inference(nnf_transformation, [], [f22])).
fof(f2785, plain, (~ spl140_284 | spl140_95), inference(avatar_split_clause, [], [f752, f1632, f2781])).
fof(f752, plain, ((e4 = op(e1, e1)) | ~ sP13), inference(cnf_transformation, [], [f277])).
fof(f277, plain, ((~ (e1 = op(e1, e4)) & (e4 = op(e1, e1))) | ~ sP13), inference(nnf_transformation, [], [f23])).
fof(f2778, plain, (~ spl140_283 | ~ spl140_73), inference(avatar_split_clause, [], [f1162, f1540, f2775])).
fof(f1162, plain, (~ (e2 = op(e2, unit)) | ~ sP14), inference(definition_unfolding, [], [f751, f377])).
fof(f751, plain, (~ (e2 = op(e2, e0)) | ~ sP14), inference(cnf_transformation, [], [f276])).
fof(f276, plain, ((~ (e2 = op(e2, e0)) & (e0 = op(e2, e2))) | ~ sP14), inference(nnf_transformation, [], [f24])).
fof(f2773, plain, (~ spl140_282 | spl140_62), inference(avatar_split_clause, [], [f748, f1494, f2769])).
fof(f748, plain, ((e1 = op(e2, e2)) | ~ sP15), inference(cnf_transformation, [], [f275])).
fof(f275, plain, ((~ (e2 = op(e2, e1)) & (e1 = op(e2, e2))) | ~ sP15), inference(nnf_transformation, [], [f25])).
fof(f2767, plain, (~ spl140_281 | spl140_63), inference(avatar_split_clause, [], [f746, f1498, f2763])).
fof(f746, plain, ((e2 = op(e2, e2)) | ~ sP16), inference(cnf_transformation, [], [f274])).
fof(f274, plain, ((~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP16), inference(nnf_transformation, [], [f26])).
fof(f2761, plain, (~ spl140_280 | spl140_64), inference(avatar_split_clause, [], [f744, f1502, f2757])).
fof(f744, plain, ((e3 = op(e2, e2)) | ~ sP17), inference(cnf_transformation, [], [f273])).
fof(f273, plain, ((~ (e2 = op(e2, e3)) & (e3 = op(e2, e2))) | ~ sP17), inference(nnf_transformation, [], [f27])).
fof(f2755, plain, (~ spl140_279 | spl140_65), inference(avatar_split_clause, [], [f742, f1506, f2751])).
fof(f742, plain, ((e4 = op(e2, e2)) | ~ sP18), inference(cnf_transformation, [], [f272])).
fof(f272, plain, ((~ (e2 = op(e2, e4)) & (e4 = op(e2, e2))) | ~ sP18), inference(nnf_transformation, [], [f28])).
fof(f2748, plain, (~ spl140_278 | ~ spl140_49), inference(avatar_split_clause, [], [f1160, f1439, f2745])).
fof(f1160, plain, (~ (e3 = op(e3, unit)) | ~ sP19), inference(definition_unfolding, [], [f741, f377])).
fof(f741, plain, (~ (e3 = op(e3, e0)) | ~ sP19), inference(cnf_transformation, [], [f271])).
fof(f271, plain, ((~ (e3 = op(e3, e0)) & (e0 = op(e3, e3))) | ~ sP19), inference(nnf_transformation, [], [f29])).
fof(f2743, plain, (~ spl140_277 | spl140_32), inference(avatar_split_clause, [], [f738, f1368, f2739])).
fof(f738, plain, ((e1 = op(e3, e3)) | ~ sP20), inference(cnf_transformation, [], [f270])).
fof(f270, plain, ((~ (e3 = op(e3, e1)) & (e1 = op(e3, e3))) | ~ sP20), inference(nnf_transformation, [], [f30])).
fof(f2737, plain, (~ spl140_276 | spl140_33), inference(avatar_split_clause, [], [f736, f1372, f2733])).
fof(f736, plain, ((e2 = op(e3, e3)) | ~ sP21), inference(cnf_transformation, [], [f269])).
fof(f269, plain, ((~ (e3 = op(e3, e2)) & (e2 = op(e3, e3))) | ~ sP21), inference(nnf_transformation, [], [f31])).
fof(f2731, plain, (~ spl140_275 | spl140_34), inference(avatar_split_clause, [], [f734, f1376, f2727])).
fof(f734, plain, ((e3 = op(e3, e3)) | ~ sP22), inference(cnf_transformation, [], [f268])).
fof(f268, plain, ((~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ~ sP22), inference(nnf_transformation, [], [f32])).
fof(f2725, plain, (~ spl140_274 | spl140_35), inference(avatar_split_clause, [], [f732, f1380, f2721])).
fof(f732, plain, ((e4 = op(e3, e3)) | ~ sP23), inference(cnf_transformation, [], [f267])).
fof(f267, plain, ((~ (e3 = op(e3, e4)) & (e4 = op(e3, e3))) | ~ sP23), inference(nnf_transformation, [], [f33])).
fof(f2718, plain, (~ spl140_273 | ~ spl140_25), inference(avatar_split_clause, [], [f1158, f1338, f2715])).
fof(f1158, plain, (~ (e4 = op(e4, unit)) | ~ sP24), inference(definition_unfolding, [], [f731, f377])).
fof(f731, plain, (~ (e4 = op(e4, e0)) | ~ sP24), inference(cnf_transformation, [], [f266])).
fof(f266, plain, ((~ (e4 = op(e4, e0)) & (e0 = op(e4, e4))) | ~ sP24), inference(nnf_transformation, [], [f34])).
fof(f2713, plain, (~ spl140_272 | spl140_2), inference(avatar_split_clause, [], [f728, f1242, f2709])).
fof(f728, plain, ((e1 = op(e4, e4)) | ~ sP25), inference(cnf_transformation, [], [f265])).
fof(f265, plain, ((~ (e4 = op(e4, e1)) & (e1 = op(e4, e4))) | ~ sP25), inference(nnf_transformation, [], [f35])).
fof(f2707, plain, (~ spl140_263 | spl140_270 | ~ spl140_271), inference(avatar_split_clause, [], [f1156, f2704, f2700, f2669])).
fof(f1156, plain, (~ (e1 = op(unit, op(e1, unit))) | (unit = op(e1, op(e1, unit))) | ~ sP26), inference(definition_unfolding, [], [f724, f377, f377, f377, f377])).
fof(f724, plain, (~ (e1 = op(e0, op(e1, e0))) | (e0 = op(e1, op(e1, e0))) | ~ sP26), inference(cnf_transformation, [], [f264])).
fof(f264, plain, (((~ (e4 = op(e0, op(e4, e0))) | (e0 = op(e4, op(e4, e0)))) & (~ (e3 = op(e0, op(e3, e0))) | (e0 = op(e3, op(e3, e0)))) & (~ (e2 = op(e0, op(e2, e0))) | (e0 = op(e2, op(e2, e0)))) & (~ (e1 = op(e0, op(e1, e0))) | (e0 = op(e1, op(e1, e0)))) & (~ (e0 = op(e0, op(e0, e0))) | (e0 = op(e0, op(e0, e0))))) | ~ sP26), inference(nnf_transformation, [], [f36])).
fof(f2689, plain, (~ spl140_263 | spl140_266 | ~ spl140_267), inference(avatar_split_clause, [], [f1154, f2686, f2682, f2669])).
fof(f1154, plain, (~ (e3 = op(unit, op(e3, unit))) | (unit = op(e3, op(e3, unit))) | ~ sP26), inference(definition_unfolding, [], [f726, f377, f377, f377, f377])).
fof(f726, plain, (~ (e3 = op(e0, op(e3, e0))) | (e0 = op(e3, op(e3, e0))) | ~ sP26), inference(cnf_transformation, [], [f264])).
fof(f2680, plain, (~ spl140_263 | spl140_264 | ~ spl140_265), inference(avatar_split_clause, [], [f1153, f2677, f2673, f2669])).
fof(f1153, plain, (~ (e4 = op(unit, op(e4, unit))) | (unit = op(e4, op(e4, unit))) | ~ sP26), inference(definition_unfolding, [], [f727, f377, f377, f377, f377])).
fof(f727, plain, (~ (e4 = op(e0, op(e4, e0))) | (e0 = op(e4, op(e4, e0))) | ~ sP26), inference(cnf_transformation, [], [f264])).
fof(f2658, plain, (~ spl140_254 | spl140_259 | ~ spl140_260), inference(avatar_split_clause, [], [f720, f2655, f2651, f2629])).
fof(f720, plain, (~ (e2 = op(e1, op(e2, e1))) | (e1 = op(e2, op(e2, e1))) | ~ sP27), inference(cnf_transformation, [], [f263])).
fof(f263, plain, (((~ (e4 = op(e1, op(e4, e1))) | (e1 = op(e4, op(e4, e1)))) & (~ (e3 = op(e1, op(e3, e1))) | (e1 = op(e3, op(e3, e1)))) & (~ (e2 = op(e1, op(e2, e1))) | (e1 = op(e2, op(e2, e1)))) & (~ (e1 = op(e1, op(e1, e1))) | (e1 = op(e1, op(e1, e1)))) & (~ (e0 = op(e1, op(e0, e1))) | (e1 = op(e0, op(e0, e1))))) | ~ sP27), inference(nnf_transformation, [], [f37])).
fof(f2649, plain, (~ spl140_254 | spl140_257 | ~ spl140_258), inference(avatar_split_clause, [], [f721, f2646, f2642, f2629])).
fof(f721, plain, (~ (e3 = op(e1, op(e3, e1))) | (e1 = op(e3, op(e3, e1))) | ~ sP27), inference(cnf_transformation, [], [f263])).
fof(f2640, plain, (~ spl140_254 | spl140_255 | ~ spl140_256), inference(avatar_split_clause, [], [f722, f2637, f2633, f2629])).
fof(f722, plain, (~ (e4 = op(e1, op(e4, e1))) | (e1 = op(e4, op(e4, e1))) | ~ sP27), inference(cnf_transformation, [], [f263])).
fof(f2618, plain, (~ spl140_245 | spl140_250 | ~ spl140_251), inference(avatar_split_clause, [], [f714, f2615, f2611, f2589])).
fof(f714, plain, (~ (e1 = op(e2, op(e1, e2))) | (e2 = op(e1, op(e1, e2))) | ~ sP28), inference(cnf_transformation, [], [f262])).
fof(f262, plain, (((~ (e4 = op(e2, op(e4, e2))) | (e2 = op(e4, op(e4, e2)))) & (~ (e3 = op(e2, op(e3, e2))) | (e2 = op(e3, op(e3, e2)))) & (~ (e2 = op(e2, op(e2, e2))) | (e2 = op(e2, op(e2, e2)))) & (~ (e1 = op(e2, op(e1, e2))) | (e2 = op(e1, op(e1, e2)))) & (~ (e0 = op(e2, op(e0, e2))) | (e2 = op(e0, op(e0, e2))))) | ~ sP28), inference(nnf_transformation, [], [f38])).
fof(f2609, plain, (~ spl140_245 | spl140_248 | ~ spl140_249), inference(avatar_split_clause, [], [f716, f2606, f2602, f2589])).
fof(f716, plain, (~ (e3 = op(e2, op(e3, e2))) | (e2 = op(e3, op(e3, e2))) | ~ sP28), inference(cnf_transformation, [], [f262])).
fof(f2600, plain, (~ spl140_245 | spl140_246 | ~ spl140_247), inference(avatar_split_clause, [], [f717, f2597, f2593, f2589])).
fof(f717, plain, (~ (e4 = op(e2, op(e4, e2))) | (e2 = op(e4, op(e4, e2))) | ~ sP28), inference(cnf_transformation, [], [f262])).
fof(f2578, plain, (~ spl140_236 | spl140_241 | ~ spl140_242), inference(avatar_split_clause, [], [f709, f2575, f2571, f2549])).
fof(f709, plain, (~ (e1 = op(e3, op(e1, e3))) | (e3 = op(e1, op(e1, e3))) | ~ sP29), inference(cnf_transformation, [], [f261])).
fof(f261, plain, (((~ (e4 = op(e3, op(e4, e3))) | (e3 = op(e4, op(e4, e3)))) & (~ (e3 = op(e3, op(e3, e3))) | (e3 = op(e3, op(e3, e3)))) & (~ (e2 = op(e3, op(e2, e3))) | (e3 = op(e2, op(e2, e3)))) & (~ (e1 = op(e3, op(e1, e3))) | (e3 = op(e1, op(e1, e3)))) & (~ (e0 = op(e3, op(e0, e3))) | (e3 = op(e0, op(e0, e3))))) | ~ sP29), inference(nnf_transformation, [], [f39])).
fof(f2569, plain, (~ spl140_236 | spl140_239 | ~ spl140_240), inference(avatar_split_clause, [], [f710, f2566, f2562, f2549])).
fof(f710, plain, (~ (e2 = op(e3, op(e2, e3))) | (e3 = op(e2, op(e2, e3))) | ~ sP29), inference(cnf_transformation, [], [f261])).
fof(f2560, plain, (~ spl140_236 | spl140_237 | ~ spl140_238), inference(avatar_split_clause, [], [f712, f2557, f2553, f2549])).
fof(f712, plain, (~ (e4 = op(e3, op(e4, e3))) | (e3 = op(e4, op(e4, e3))) | ~ sP29), inference(cnf_transformation, [], [f261])).
fof(f2437, plain, (~ spl140_213 | spl140_122), inference(avatar_split_clause, [], [f1105, f1746, f2433])).
fof(f1105, plain, ((e1 = op(unit, unit)) | ~ sP52), inference(definition_unfolding, [], [f662, f377, f377])).
fof(f662, plain, ((op(e0, e0) = e1) | ~ sP52), inference(cnf_transformation, [], [f238])).
fof(f238, plain, ((((e1 = unit) | ~ (e0 = op(e0, e1))) & (op(e0, e0) = e1)) | ~ sP52), inference(nnf_transformation, [], [f62])).
fof(f2429, plain, (~ spl140_212 | ~ spl140_117), inference(avatar_split_clause, [], [f2424, f1725, f2426])).
fof(f2424, plain, (~ (e1 = op(unit, e1)) | ~ sP53), inference(subsumption_resolution, [], [f1102, f1033])).
fof(f1102, plain, ((e1 = unit) | ~ (e1 = op(unit, e1)) | ~ sP53), inference(definition_unfolding, [], [f661, f377])).
fof(f661, plain, ((e1 = unit) | ~ (e1 = op(e0, e1)) | ~ sP53), inference(cnf_transformation, [], [f237])).
fof(f237, plain, ((((e1 = unit) | ~ (e1 = op(e0, e1))) & (e1 = op(e0, e1))) | ~ sP53), inference(nnf_transformation, [], [f63])).
fof(f2423, plain, (~ spl140_211 | spl140_112), inference(avatar_split_clause, [], [f1101, f1704, f2419])).
fof(f1101, plain, ((e1 = op(unit, e2)) | ~ sP54), inference(definition_unfolding, [], [f658, f377])).
fof(f658, plain, ((e1 = op(e0, e2)) | ~ sP54), inference(cnf_transformation, [], [f236])).
fof(f236, plain, ((((e1 = unit) | ~ (e2 = op(e0, e1))) & (e1 = op(e0, e2))) | ~ sP54), inference(nnf_transformation, [], [f64])).
fof(f2416, plain, (~ spl140_210 | spl140_107), inference(avatar_split_clause, [], [f1099, f1683, f2412])).
fof(f1099, plain, ((e1 = op(unit, e3)) | ~ sP55), inference(definition_unfolding, [], [f656, f377])).
fof(f656, plain, ((e1 = op(e0, e3)) | ~ sP55), inference(cnf_transformation, [], [f235])).
fof(f235, plain, ((((e1 = unit) | ~ (e3 = op(e0, e1))) & (e1 = op(e0, e3))) | ~ sP55), inference(nnf_transformation, [], [f65])).
fof(f2409, plain, (~ spl140_209 | spl140_102), inference(avatar_split_clause, [], [f1097, f1662, f2405])).
fof(f1097, plain, ((e1 = op(unit, e4)) | ~ sP56), inference(definition_unfolding, [], [f654, f377])).
fof(f654, plain, ((e1 = op(e0, e4)) | ~ sP56), inference(cnf_transformation, [], [f234])).
fof(f234, plain, ((((e1 = unit) | ~ (e4 = op(e0, e1))) & (e1 = op(e0, e4))) | ~ sP56), inference(nnf_transformation, [], [f66])).
fof(f2401, plain, (~ spl140_208 | ~ spl140_91), inference(avatar_split_clause, [], [f2396, f1616, f2398])).
fof(f2396, plain, (~ (op(e1, e1) = unit) | ~ sP57), inference(subsumption_resolution, [], [f1094, f1033])).
fof(f1094, plain, ((e1 = unit) | ~ (op(e1, e1) = unit) | ~ sP57), inference(definition_unfolding, [], [f653, f377])).
fof(f653, plain, ((e1 = unit) | ~ (e0 = op(e1, e1)) | ~ sP57), inference(cnf_transformation, [], [f233])).
fof(f233, plain, ((((e1 = unit) | ~ (e0 = op(e1, e1))) & (e1 = op(e1, e0))) | ~ sP57), inference(nnf_transformation, [], [f67])).
fof(f2395, plain, (~ spl140_207 | spl140_92), inference(avatar_split_clause, [], [f650, f1620, f2391])).
fof(f650, plain, ((e1 = op(e1, e1)) | ~ sP58), inference(cnf_transformation, [], [f232])).
fof(f232, plain, ((((e1 = unit) | ~ (e1 = op(e1, e1))) & (e1 = op(e1, e1))) | ~ sP58), inference(nnf_transformation, [], [f68])).
fof(f2388, plain, (~ spl140_206 | spl140_87), inference(avatar_split_clause, [], [f648, f1599, f2384])).
fof(f648, plain, ((e1 = op(e1, e2)) | ~ sP59), inference(cnf_transformation, [], [f231])).
fof(f231, plain, ((((e1 = unit) | ~ (e2 = op(e1, e1))) & (e1 = op(e1, e2))) | ~ sP59), inference(nnf_transformation, [], [f69])).
fof(f2381, plain, (~ spl140_205 | spl140_82), inference(avatar_split_clause, [], [f646, f1578, f2377])).
fof(f646, plain, ((e1 = op(e1, e3)) | ~ sP60), inference(cnf_transformation, [], [f230])).
fof(f230, plain, ((((e1 = unit) | ~ (e3 = op(e1, e1))) & (e1 = op(e1, e3))) | ~ sP60), inference(nnf_transformation, [], [f70])).
fof(f2374, plain, (~ spl140_204 | spl140_77), inference(avatar_split_clause, [], [f644, f1557, f2370])).
fof(f644, plain, ((e1 = op(e1, e4)) | ~ sP61), inference(cnf_transformation, [], [f229])).
fof(f229, plain, ((((e1 = unit) | ~ (e4 = op(e1, e1))) & (e1 = op(e1, e4))) | ~ sP61), inference(nnf_transformation, [], [f71])).
fof(f2367, plain, (~ spl140_203 | spl140_72), inference(avatar_split_clause, [], [f1093, f1536, f2363])).
fof(f1093, plain, ((e1 = op(e2, unit)) | ~ sP62), inference(definition_unfolding, [], [f642, f377])).
fof(f642, plain, ((e1 = op(e2, e0)) | ~ sP62), inference(cnf_transformation, [], [f228])).
fof(f228, plain, ((((e1 = unit) | ~ (e0 = op(e2, e1))) & (e1 = op(e2, e0))) | ~ sP62), inference(nnf_transformation, [], [f72])).
fof(f2360, plain, (~ spl140_202 | spl140_67), inference(avatar_split_clause, [], [f640, f1515, f2356])).
fof(f640, plain, ((e1 = op(e2, e1)) | ~ sP63), inference(cnf_transformation, [], [f227])).
fof(f227, plain, ((((e1 = unit) | ~ (e1 = op(e2, e1))) & (e1 = op(e2, e1))) | ~ sP63), inference(nnf_transformation, [], [f73])).
fof(f2353, plain, (~ spl140_201 | spl140_62), inference(avatar_split_clause, [], [f638, f1494, f2349])).
fof(f638, plain, ((e1 = op(e2, e2)) | ~ sP64), inference(cnf_transformation, [], [f226])).
fof(f226, plain, ((((e1 = unit) | ~ (e2 = op(e2, e1))) & (e1 = op(e2, e2))) | ~ sP64), inference(nnf_transformation, [], [f74])).
fof(f2346, plain, (~ spl140_200 | spl140_57), inference(avatar_split_clause, [], [f636, f1473, f2342])).
fof(f636, plain, ((e1 = op(e2, e3)) | ~ sP65), inference(cnf_transformation, [], [f225])).
fof(f225, plain, ((((e1 = unit) | ~ (e3 = op(e2, e1))) & (e1 = op(e2, e3))) | ~ sP65), inference(nnf_transformation, [], [f75])).
fof(f2345, plain, (~ spl140_200 | ~ spl140_69), inference(avatar_split_clause, [], [f2340, f1523, f2342])).
fof(f2340, plain, (~ (e3 = op(e2, e1)) | ~ sP65), inference(subsumption_resolution, [], [f637, f1033])).
fof(f637, plain, ((e1 = unit) | ~ (e3 = op(e2, e1)) | ~ sP65), inference(cnf_transformation, [], [f225])).
fof(f2339, plain, (~ spl140_199 | spl140_52), inference(avatar_split_clause, [], [f634, f1452, f2335])).
fof(f634, plain, ((e1 = op(e2, e4)) | ~ sP66), inference(cnf_transformation, [], [f224])).
fof(f224, plain, ((((e1 = unit) | ~ (e4 = op(e2, e1))) & (e1 = op(e2, e4))) | ~ sP66), inference(nnf_transformation, [], [f76])).
fof(f2338, plain, (~ spl140_199 | ~ spl140_70), inference(avatar_split_clause, [], [f2333, f1527, f2335])).
fof(f2333, plain, (~ (e4 = op(e2, e1)) | ~ sP66), inference(subsumption_resolution, [], [f635, f1033])).
fof(f635, plain, ((e1 = unit) | ~ (e4 = op(e2, e1)) | ~ sP66), inference(cnf_transformation, [], [f224])).
fof(f2332, plain, (~ spl140_198 | spl140_47), inference(avatar_split_clause, [], [f1091, f1431, f2328])).
fof(f1091, plain, ((e1 = op(e3, unit)) | ~ sP67), inference(definition_unfolding, [], [f632, f377])).
fof(f632, plain, ((e1 = op(e3, e0)) | ~ sP67), inference(cnf_transformation, [], [f223])).
fof(f223, plain, ((((e1 = unit) | ~ (e0 = op(e3, e1))) & (e1 = op(e3, e0))) | ~ sP67), inference(nnf_transformation, [], [f77])).
fof(f2325, plain, (~ spl140_197 | spl140_42), inference(avatar_split_clause, [], [f630, f1410, f2321])).
fof(f630, plain, ((e1 = op(e3, e1)) | ~ sP68), inference(cnf_transformation, [], [f222])).
fof(f222, plain, ((((e1 = unit) | ~ (e1 = op(e3, e1))) & (e1 = op(e3, e1))) | ~ sP68), inference(nnf_transformation, [], [f78])).
fof(f2318, plain, (~ spl140_196 | spl140_37), inference(avatar_split_clause, [], [f628, f1389, f2314])).
fof(f628, plain, ((e1 = op(e3, e2)) | ~ sP69), inference(cnf_transformation, [], [f221])).
fof(f221, plain, ((((e1 = unit) | ~ (e2 = op(e3, e1))) & (e1 = op(e3, e2))) | ~ sP69), inference(nnf_transformation, [], [f79])).
fof(f2317, plain, (~ spl140_196 | ~ spl140_43), inference(avatar_split_clause, [], [f2312, f1414, f2314])).
fof(f2312, plain, (~ (e2 = op(e3, e1)) | ~ sP69), inference(subsumption_resolution, [], [f629, f1033])).
fof(f629, plain, ((e1 = unit) | ~ (e2 = op(e3, e1)) | ~ sP69), inference(cnf_transformation, [], [f221])).
fof(f2311, plain, (~ spl140_195 | spl140_32), inference(avatar_split_clause, [], [f626, f1368, f2307])).
fof(f626, plain, ((e1 = op(e3, e3)) | ~ sP70), inference(cnf_transformation, [], [f220])).
fof(f220, plain, ((((e1 = unit) | ~ (e3 = op(e3, e1))) & (e1 = op(e3, e3))) | ~ sP70), inference(nnf_transformation, [], [f80])).
fof(f2304, plain, (~ spl140_194 | spl140_27), inference(avatar_split_clause, [], [f624, f1347, f2300])).
fof(f624, plain, ((e1 = op(e3, e4)) | ~ sP71), inference(cnf_transformation, [], [f219])).
fof(f219, plain, ((((e1 = unit) | ~ (e4 = op(e3, e1))) & (e1 = op(e3, e4))) | ~ sP71), inference(nnf_transformation, [], [f81])).
fof(f2303, plain, (~ spl140_194 | ~ spl140_45), inference(avatar_split_clause, [], [f2298, f1422, f2300])).
fof(f2298, plain, (~ (e4 = op(e3, e1)) | ~ sP71), inference(subsumption_resolution, [], [f625, f1033])).
fof(f625, plain, ((e1 = unit) | ~ (e4 = op(e3, e1)) | ~ sP71), inference(cnf_transformation, [], [f219])).
fof(f2297, plain, (~ spl140_193 | spl140_22), inference(avatar_split_clause, [], [f1089, f1326, f2293])).
fof(f1089, plain, ((e1 = op(e4, unit)) | ~ sP72), inference(definition_unfolding, [], [f622, f377])).
fof(f622, plain, ((e1 = op(e4, e0)) | ~ sP72), inference(cnf_transformation, [], [f218])).
fof(f218, plain, ((((e1 = unit) | ~ (e0 = op(e4, e1))) & (e1 = op(e4, e0))) | ~ sP72), inference(nnf_transformation, [], [f82])).
fof(f2290, plain, (~ spl140_192 | spl140_17), inference(avatar_split_clause, [], [f620, f1305, f2286])).
fof(f620, plain, ((e1 = op(e4, e1)) | ~ sP73), inference(cnf_transformation, [], [f217])).
fof(f217, plain, ((((e1 = unit) | ~ (e1 = op(e4, e1))) & (e1 = op(e4, e1))) | ~ sP73), inference(nnf_transformation, [], [f83])).
fof(f2283, plain, (~ spl140_191 | spl140_123), inference(avatar_split_clause, [], [f1087, f1750, f2279])).
fof(f1087, plain, ((e2 = op(unit, unit)) | ~ sP74), inference(definition_unfolding, [], [f618, f377, f377])).
fof(f618, plain, ((op(e0, e0) = e2) | ~ sP74), inference(cnf_transformation, [], [f216])).
fof(f216, plain, ((((e2 = unit) | ~ (e0 = op(e0, e2))) & (op(e0, e0) = e2)) | ~ sP74), inference(nnf_transformation, [], [f84])).
fof(f2276, plain, (~ spl140_190 | spl140_118), inference(avatar_split_clause, [], [f1085, f1729, f2272])).
fof(f1085, plain, ((e2 = op(unit, e1)) | ~ sP75), inference(definition_unfolding, [], [f616, f377])).
fof(f616, plain, ((e2 = op(e0, e1)) | ~ sP75), inference(cnf_transformation, [], [f215])).
fof(f215, plain, ((((e2 = unit) | ~ (e1 = op(e0, e2))) & (e2 = op(e0, e1))) | ~ sP75), inference(nnf_transformation, [], [f85])).
fof(f2268, plain, (~ spl140_189 | ~ spl140_113), inference(avatar_split_clause, [], [f2263, f1708, f2265])).
fof(f2263, plain, (~ (e2 = op(unit, e2)) | ~ sP76), inference(subsumption_resolution, [], [f1082, f1032])).
fof(f1082, plain, ((e2 = unit) | ~ (e2 = op(unit, e2)) | ~ sP76), inference(definition_unfolding, [], [f615, f377])).
fof(f615, plain, ((e2 = unit) | ~ (e2 = op(e0, e2)) | ~ sP76), inference(cnf_transformation, [], [f214])).
fof(f214, plain, ((((e2 = unit) | ~ (e2 = op(e0, e2))) & (e2 = op(e0, e2))) | ~ sP76), inference(nnf_transformation, [], [f86])).
fof(f2262, plain, (~ spl140_188 | spl140_108), inference(avatar_split_clause, [], [f1081, f1687, f2258])).
fof(f1081, plain, ((e2 = op(unit, e3)) | ~ sP77), inference(definition_unfolding, [], [f612, f377])).
fof(f612, plain, ((e2 = op(e0, e3)) | ~ sP77), inference(cnf_transformation, [], [f213])).
fof(f213, plain, ((((e2 = unit) | ~ (e3 = op(e0, e2))) & (e2 = op(e0, e3))) | ~ sP77), inference(nnf_transformation, [], [f87])).
fof(f2255, plain, (~ spl140_187 | spl140_103), inference(avatar_split_clause, [], [f1079, f1666, f2251])).
fof(f1079, plain, ((e2 = op(unit, e4)) | ~ sP78), inference(definition_unfolding, [], [f610, f377])).
fof(f610, plain, ((e2 = op(e0, e4)) | ~ sP78), inference(cnf_transformation, [], [f212])).
fof(f212, plain, ((((e2 = unit) | ~ (e4 = op(e0, e2))) & (e2 = op(e0, e4))) | ~ sP78), inference(nnf_transformation, [], [f88])).
fof(f2248, plain, (~ spl140_186 | spl140_98), inference(avatar_split_clause, [], [f1077, f1645, f2244])).
fof(f1077, plain, ((e2 = op(e1, unit)) | ~ sP79), inference(definition_unfolding, [], [f608, f377])).
fof(f608, plain, ((e2 = op(e1, e0)) | ~ sP79), inference(cnf_transformation, [], [f211])).
fof(f211, plain, ((((e2 = unit) | ~ (e0 = op(e1, e2))) & (e2 = op(e1, e0))) | ~ sP79), inference(nnf_transformation, [], [f89])).
fof(f2241, plain, (~ spl140_185 | spl140_93), inference(avatar_split_clause, [], [f606, f1624, f2237])).
fof(f606, plain, ((e2 = op(e1, e1)) | ~ sP80), inference(cnf_transformation, [], [f210])).
fof(f210, plain, ((((e2 = unit) | ~ (e1 = op(e1, e2))) & (e2 = op(e1, e1))) | ~ sP80), inference(nnf_transformation, [], [f90])).
fof(f2234, plain, (~ spl140_184 | spl140_88), inference(avatar_split_clause, [], [f604, f1603, f2230])).
fof(f604, plain, ((e2 = op(e1, e2)) | ~ sP81), inference(cnf_transformation, [], [f209])).
fof(f209, plain, ((((e2 = unit) | ~ (e2 = op(e1, e2))) & (e2 = op(e1, e2))) | ~ sP81), inference(nnf_transformation, [], [f91])).
fof(f2227, plain, (~ spl140_183 | spl140_83), inference(avatar_split_clause, [], [f602, f1582, f2223])).
fof(f602, plain, ((e2 = op(e1, e3)) | ~ sP82), inference(cnf_transformation, [], [f208])).
fof(f208, plain, ((((e2 = unit) | ~ (e3 = op(e1, e2))) & (e2 = op(e1, e3))) | ~ sP82), inference(nnf_transformation, [], [f92])).
fof(f2226, plain, (~ spl140_183 | ~ spl140_89), inference(avatar_split_clause, [], [f2221, f1607, f2223])).
fof(f2221, plain, (~ (e3 = op(e1, e2)) | ~ sP82), inference(subsumption_resolution, [], [f603, f1032])).
fof(f603, plain, ((e2 = unit) | ~ (e3 = op(e1, e2)) | ~ sP82), inference(cnf_transformation, [], [f208])).
fof(f2220, plain, (~ spl140_182 | spl140_78), inference(avatar_split_clause, [], [f600, f1561, f2216])).
fof(f600, plain, ((e2 = op(e1, e4)) | ~ sP83), inference(cnf_transformation, [], [f207])).
fof(f207, plain, ((((e2 = unit) | ~ (e4 = op(e1, e2))) & (e2 = op(e1, e4))) | ~ sP83), inference(nnf_transformation, [], [f93])).
fof(f2219, plain, (~ spl140_182 | ~ spl140_90), inference(avatar_split_clause, [], [f2214, f1611, f2216])).
fof(f2214, plain, (~ (e4 = op(e1, e2)) | ~ sP83), inference(subsumption_resolution, [], [f601, f1032])).
fof(f601, plain, ((e2 = unit) | ~ (e4 = op(e1, e2)) | ~ sP83), inference(cnf_transformation, [], [f207])).
fof(f2212, plain, (~ spl140_181 | ~ spl140_61), inference(avatar_split_clause, [], [f2207, f1490, f2209])).
fof(f2207, plain, (~ (op(e2, e2) = unit) | ~ sP84), inference(subsumption_resolution, [], [f1074, f1032])).
fof(f1074, plain, ((e2 = unit) | ~ (op(e2, e2) = unit) | ~ sP84), inference(definition_unfolding, [], [f599, f377])).
fof(f599, plain, ((e2 = unit) | ~ (e0 = op(e2, e2)) | ~ sP84), inference(cnf_transformation, [], [f206])).
fof(f206, plain, ((((e2 = unit) | ~ (e0 = op(e2, e2))) & (e2 = op(e2, e0))) | ~ sP84), inference(nnf_transformation, [], [f94])).
fof(f2206, plain, (~ spl140_180 | spl140_68), inference(avatar_split_clause, [], [f596, f1519, f2202])).
fof(f596, plain, ((e2 = op(e2, e1)) | ~ sP85), inference(cnf_transformation, [], [f205])).
fof(f205, plain, ((((e2 = unit) | ~ (e1 = op(e2, e2))) & (e2 = op(e2, e1))) | ~ sP85), inference(nnf_transformation, [], [f95])).
fof(f2199, plain, (~ spl140_179 | spl140_63), inference(avatar_split_clause, [], [f594, f1498, f2195])).
fof(f594, plain, ((e2 = op(e2, e2)) | ~ sP86), inference(cnf_transformation, [], [f204])).
fof(f204, plain, ((((e2 = unit) | ~ (e2 = op(e2, e2))) & (e2 = op(e2, e2))) | ~ sP86), inference(nnf_transformation, [], [f96])).
fof(f2192, plain, (~ spl140_178 | spl140_58), inference(avatar_split_clause, [], [f592, f1477, f2188])).
fof(f592, plain, ((e2 = op(e2, e3)) | ~ sP87), inference(cnf_transformation, [], [f203])).
fof(f203, plain, ((((e2 = unit) | ~ (e3 = op(e2, e2))) & (e2 = op(e2, e3))) | ~ sP87), inference(nnf_transformation, [], [f97])).
fof(f2185, plain, (~ spl140_177 | spl140_53), inference(avatar_split_clause, [], [f590, f1456, f2181])).
fof(f590, plain, ((e2 = op(e2, e4)) | ~ sP88), inference(cnf_transformation, [], [f202])).
fof(f202, plain, ((((e2 = unit) | ~ (e4 = op(e2, e2))) & (e2 = op(e2, e4))) | ~ sP88), inference(nnf_transformation, [], [f98])).
fof(f2178, plain, (~ spl140_176 | spl140_48), inference(avatar_split_clause, [], [f1073, f1435, f2174])).
fof(f1073, plain, ((e2 = op(e3, unit)) | ~ sP89), inference(definition_unfolding, [], [f588, f377])).
fof(f588, plain, ((e2 = op(e3, e0)) | ~ sP89), inference(cnf_transformation, [], [f201])).
fof(f201, plain, ((((e2 = unit) | ~ (e0 = op(e3, e2))) & (e2 = op(e3, e0))) | ~ sP89), inference(nnf_transformation, [], [f99])).
fof(f2171, plain, (~ spl140_175 | spl140_43), inference(avatar_split_clause, [], [f586, f1414, f2167])).
fof(f586, plain, ((e2 = op(e3, e1)) | ~ sP90), inference(cnf_transformation, [], [f200])).
fof(f200, plain, ((((e2 = unit) | ~ (e1 = op(e3, e2))) & (e2 = op(e3, e1))) | ~ sP90), inference(nnf_transformation, [], [f100])).
fof(f2170, plain, (~ spl140_175 | ~ spl140_37), inference(avatar_split_clause, [], [f2165, f1389, f2167])).
fof(f2165, plain, (~ (e1 = op(e3, e2)) | ~ sP90), inference(subsumption_resolution, [], [f587, f1032])).
fof(f587, plain, ((e2 = unit) | ~ (e1 = op(e3, e2)) | ~ sP90), inference(cnf_transformation, [], [f200])).
fof(f2164, plain, (~ spl140_174 | spl140_38), inference(avatar_split_clause, [], [f584, f1393, f2160])).
fof(f584, plain, ((e2 = op(e3, e2)) | ~ sP91), inference(cnf_transformation, [], [f199])).
fof(f199, plain, ((((e2 = unit) | ~ (e2 = op(e3, e2))) & (e2 = op(e3, e2))) | ~ sP91), inference(nnf_transformation, [], [f101])).
fof(f2157, plain, (~ spl140_173 | spl140_33), inference(avatar_split_clause, [], [f582, f1372, f2153])).
fof(f582, plain, ((e2 = op(e3, e3)) | ~ sP92), inference(cnf_transformation, [], [f198])).
fof(f198, plain, ((((e2 = unit) | ~ (e3 = op(e3, e2))) & (e2 = op(e3, e3))) | ~ sP92), inference(nnf_transformation, [], [f102])).
fof(f2150, plain, (~ spl140_172 | spl140_28), inference(avatar_split_clause, [], [f580, f1351, f2146])).
fof(f580, plain, ((e2 = op(e3, e4)) | ~ sP93), inference(cnf_transformation, [], [f197])).
fof(f197, plain, ((((e2 = unit) | ~ (e4 = op(e3, e2))) & (e2 = op(e3, e4))) | ~ sP93), inference(nnf_transformation, [], [f103])).
fof(f2149, plain, (~ spl140_172 | ~ spl140_40), inference(avatar_split_clause, [], [f2144, f1401, f2146])).
fof(f2144, plain, (~ (e4 = op(e3, e2)) | ~ sP93), inference(subsumption_resolution, [], [f581, f1032])).
fof(f581, plain, ((e2 = unit) | ~ (e4 = op(e3, e2)) | ~ sP93), inference(cnf_transformation, [], [f197])).
fof(f2143, plain, (~ spl140_171 | spl140_23), inference(avatar_split_clause, [], [f1071, f1330, f2139])).
fof(f1071, plain, ((e2 = op(e4, unit)) | ~ sP94), inference(definition_unfolding, [], [f578, f377])).
fof(f578, plain, ((e2 = op(e4, e0)) | ~ sP94), inference(cnf_transformation, [], [f196])).
fof(f196, plain, ((((e2 = unit) | ~ (e0 = op(e4, e2))) & (e2 = op(e4, e0))) | ~ sP94), inference(nnf_transformation, [], [f104])).
fof(f2136, plain, (~ spl140_170 | spl140_18), inference(avatar_split_clause, [], [f576, f1309, f2132])).
fof(f576, plain, ((e2 = op(e4, e1)) | ~ sP95), inference(cnf_transformation, [], [f195])).
fof(f195, plain, ((((e2 = unit) | ~ (e1 = op(e4, e2))) & (e2 = op(e4, e1))) | ~ sP95), inference(nnf_transformation, [], [f105])).
fof(f2135, plain, (~ spl140_170 | ~ spl140_12), inference(avatar_split_clause, [], [f2130, f1284, f2132])).
fof(f2130, plain, (~ (e1 = op(e4, e2)) | ~ sP95), inference(subsumption_resolution, [], [f577, f1032])).
fof(f577, plain, ((e2 = unit) | ~ (e1 = op(e4, e2)) | ~ sP95), inference(cnf_transformation, [], [f195])).
fof(f2129, plain, (~ spl140_169 | spl140_124), inference(avatar_split_clause, [], [f1069, f1754, f2125])).
fof(f1069, plain, ((e3 = op(unit, unit)) | ~ sP96), inference(definition_unfolding, [], [f574, f377, f377])).
fof(f574, plain, ((op(e0, e0) = e3) | ~ sP96), inference(cnf_transformation, [], [f194])).
fof(f194, plain, ((((e3 = unit) | ~ (e0 = op(e0, e3))) & (op(e0, e0) = e3)) | ~ sP96), inference(nnf_transformation, [], [f106])).
fof(f2122, plain, (~ spl140_168 | spl140_119), inference(avatar_split_clause, [], [f1067, f1733, f2118])).
fof(f1067, plain, ((e3 = op(unit, e1)) | ~ sP97), inference(definition_unfolding, [], [f572, f377])).
fof(f572, plain, ((e3 = op(e0, e1)) | ~ sP97), inference(cnf_transformation, [], [f193])).
fof(f193, plain, ((((e3 = unit) | ~ (e1 = op(e0, e3))) & (e3 = op(e0, e1))) | ~ sP97), inference(nnf_transformation, [], [f107])).
fof(f2115, plain, (~ spl140_167 | spl140_114), inference(avatar_split_clause, [], [f1065, f1712, f2111])).
fof(f1065, plain, ((e3 = op(unit, e2)) | ~ sP98), inference(definition_unfolding, [], [f570, f377])).
fof(f570, plain, ((e3 = op(e0, e2)) | ~ sP98), inference(cnf_transformation, [], [f192])).
fof(f192, plain, ((((e3 = unit) | ~ (e2 = op(e0, e3))) & (e3 = op(e0, e2))) | ~ sP98), inference(nnf_transformation, [], [f108])).
fof(f2107, plain, (~ spl140_166 | ~ spl140_109), inference(avatar_split_clause, [], [f2102, f1691, f2104])).
fof(f2102, plain, (~ (e3 = op(unit, e3)) | ~ sP99), inference(subsumption_resolution, [], [f1062, f1031])).
fof(f1062, plain, ((e3 = unit) | ~ (e3 = op(unit, e3)) | ~ sP99), inference(definition_unfolding, [], [f569, f377])).
fof(f569, plain, ((e3 = unit) | ~ (e3 = op(e0, e3)) | ~ sP99), inference(cnf_transformation, [], [f191])).
fof(f191, plain, ((((e3 = unit) | ~ (e3 = op(e0, e3))) & (e3 = op(e0, e3))) | ~ sP99), inference(nnf_transformation, [], [f109])).
fof(f2101, plain, (~ spl140_165 | spl140_104), inference(avatar_split_clause, [], [f1061, f1670, f2097])).
fof(f1061, plain, ((e3 = op(unit, e4)) | ~ sP100), inference(definition_unfolding, [], [f566, f377])).
fof(f566, plain, ((e3 = op(e0, e4)) | ~ sP100), inference(cnf_transformation, [], [f190])).
fof(f190, plain, ((((e3 = unit) | ~ (e4 = op(e0, e3))) & (e3 = op(e0, e4))) | ~ sP100), inference(nnf_transformation, [], [f110])).
fof(f2094, plain, (~ spl140_164 | spl140_99), inference(avatar_split_clause, [], [f1059, f1649, f2090])).
fof(f1059, plain, ((e3 = op(e1, unit)) | ~ sP101), inference(definition_unfolding, [], [f564, f377])).
fof(f564, plain, ((e3 = op(e1, e0)) | ~ sP101), inference(cnf_transformation, [], [f189])).
fof(f189, plain, ((((e3 = unit) | ~ (e0 = op(e1, e3))) & (e3 = op(e1, e0))) | ~ sP101), inference(nnf_transformation, [], [f111])).
fof(f2087, plain, (~ spl140_163 | spl140_94), inference(avatar_split_clause, [], [f562, f1628, f2083])).
fof(f562, plain, ((e3 = op(e1, e1)) | ~ sP102), inference(cnf_transformation, [], [f188])).
fof(f188, plain, ((((e3 = unit) | ~ (e1 = op(e1, e3))) & (e3 = op(e1, e1))) | ~ sP102), inference(nnf_transformation, [], [f112])).
fof(f2080, plain, (~ spl140_162 | spl140_89), inference(avatar_split_clause, [], [f560, f1607, f2076])).
fof(f560, plain, ((e3 = op(e1, e2)) | ~ sP103), inference(cnf_transformation, [], [f187])).
fof(f187, plain, ((((e3 = unit) | ~ (e2 = op(e1, e3))) & (e3 = op(e1, e2))) | ~ sP103), inference(nnf_transformation, [], [f113])).
fof(f2079, plain, (~ spl140_162 | ~ spl140_83), inference(avatar_split_clause, [], [f2074, f1582, f2076])).
fof(f2074, plain, (~ (e2 = op(e1, e3)) | ~ sP103), inference(subsumption_resolution, [], [f561, f1031])).
fof(f561, plain, ((e3 = unit) | ~ (e2 = op(e1, e3)) | ~ sP103), inference(cnf_transformation, [], [f187])).
fof(f2073, plain, (~ spl140_161 | spl140_84), inference(avatar_split_clause, [], [f558, f1586, f2069])).
fof(f558, plain, ((e3 = op(e1, e3)) | ~ sP104), inference(cnf_transformation, [], [f186])).
fof(f186, plain, ((((e3 = unit) | ~ (e3 = op(e1, e3))) & (e3 = op(e1, e3))) | ~ sP104), inference(nnf_transformation, [], [f114])).
fof(f2066, plain, (~ spl140_160 | spl140_79), inference(avatar_split_clause, [], [f556, f1565, f2062])).
fof(f556, plain, ((e3 = op(e1, e4)) | ~ sP105), inference(cnf_transformation, [], [f185])).
fof(f185, plain, ((((e3 = unit) | ~ (e4 = op(e1, e3))) & (e3 = op(e1, e4))) | ~ sP105), inference(nnf_transformation, [], [f115])).
fof(f2065, plain, (~ spl140_160 | ~ spl140_85), inference(avatar_split_clause, [], [f2060, f1590, f2062])).
fof(f2060, plain, (~ (e4 = op(e1, e3)) | ~ sP105), inference(subsumption_resolution, [], [f557, f1031])).
fof(f557, plain, ((e3 = unit) | ~ (e4 = op(e1, e3)) | ~ sP105), inference(cnf_transformation, [], [f185])).
fof(f2059, plain, (~ spl140_159 | spl140_74), inference(avatar_split_clause, [], [f1057, f1544, f2055])).
fof(f1057, plain, ((e3 = op(e2, unit)) | ~ sP106), inference(definition_unfolding, [], [f554, f377])).
fof(f554, plain, ((e3 = op(e2, e0)) | ~ sP106), inference(cnf_transformation, [], [f184])).
fof(f184, plain, ((((e3 = unit) | ~ (e0 = op(e2, e3))) & (e3 = op(e2, e0))) | ~ sP106), inference(nnf_transformation, [], [f116])).
fof(f2052, plain, (~ spl140_158 | spl140_69), inference(avatar_split_clause, [], [f552, f1523, f2048])).
fof(f552, plain, ((e3 = op(e2, e1)) | ~ sP107), inference(cnf_transformation, [], [f183])).
fof(f183, plain, ((((e3 = unit) | ~ (e1 = op(e2, e3))) & (e3 = op(e2, e1))) | ~ sP107), inference(nnf_transformation, [], [f117])).
fof(f2051, plain, (~ spl140_158 | ~ spl140_57), inference(avatar_split_clause, [], [f2046, f1473, f2048])).
fof(f2046, plain, (~ (e1 = op(e2, e3)) | ~ sP107), inference(subsumption_resolution, [], [f553, f1031])).
fof(f553, plain, ((e3 = unit) | ~ (e1 = op(e2, e3)) | ~ sP107), inference(cnf_transformation, [], [f183])).
fof(f2045, plain, (~ spl140_157 | spl140_64), inference(avatar_split_clause, [], [f550, f1502, f2041])).
fof(f550, plain, ((e3 = op(e2, e2)) | ~ sP108), inference(cnf_transformation, [], [f182])).
fof(f182, plain, ((((e3 = unit) | ~ (e2 = op(e2, e3))) & (e3 = op(e2, e2))) | ~ sP108), inference(nnf_transformation, [], [f118])).
fof(f2038, plain, (~ spl140_156 | spl140_59), inference(avatar_split_clause, [], [f548, f1481, f2034])).
fof(f548, plain, ((e3 = op(e2, e3)) | ~ sP109), inference(cnf_transformation, [], [f181])).
fof(f181, plain, ((((e3 = unit) | ~ (e3 = op(e2, e3))) & (e3 = op(e2, e3))) | ~ sP109), inference(nnf_transformation, [], [f119])).
fof(f2031, plain, (~ spl140_155 | spl140_54), inference(avatar_split_clause, [], [f546, f1460, f2027])).
fof(f546, plain, ((e3 = op(e2, e4)) | ~ sP110), inference(cnf_transformation, [], [f180])).
fof(f180, plain, ((((e3 = unit) | ~ (e4 = op(e2, e3))) & (e3 = op(e2, e4))) | ~ sP110), inference(nnf_transformation, [], [f120])).
fof(f2030, plain, (~ spl140_155 | ~ spl140_60), inference(avatar_split_clause, [], [f2025, f1485, f2027])).
fof(f2025, plain, (~ (e4 = op(e2, e3)) | ~ sP110), inference(subsumption_resolution, [], [f547, f1031])).
fof(f547, plain, ((e3 = unit) | ~ (e4 = op(e2, e3)) | ~ sP110), inference(cnf_transformation, [], [f180])).
fof(f2023, plain, (~ spl140_154 | ~ spl140_31), inference(avatar_split_clause, [], [f2018, f1364, f2020])).
fof(f2018, plain, (~ (op(e3, e3) = unit) | ~ sP111), inference(subsumption_resolution, [], [f1054, f1031])).
fof(f1054, plain, ((e3 = unit) | ~ (op(e3, e3) = unit) | ~ sP111), inference(definition_unfolding, [], [f545, f377])).
fof(f545, plain, ((e3 = unit) | ~ (e0 = op(e3, e3)) | ~ sP111), inference(cnf_transformation, [], [f179])).
fof(f179, plain, ((((e3 = unit) | ~ (e0 = op(e3, e3))) & (e3 = op(e3, e0))) | ~ sP111), inference(nnf_transformation, [], [f121])).
fof(f2017, plain, (~ spl140_153 | spl140_44), inference(avatar_split_clause, [], [f542, f1418, f2013])).
fof(f542, plain, ((e3 = op(e3, e1)) | ~ sP112), inference(cnf_transformation, [], [f178])).
fof(f178, plain, ((((e3 = unit) | ~ (e1 = op(e3, e3))) & (e3 = op(e3, e1))) | ~ sP112), inference(nnf_transformation, [], [f122])).
fof(f2010, plain, (~ spl140_152 | spl140_39), inference(avatar_split_clause, [], [f540, f1397, f2006])).
fof(f540, plain, ((e3 = op(e3, e2)) | ~ sP113), inference(cnf_transformation, [], [f177])).
fof(f177, plain, ((((e3 = unit) | ~ (e2 = op(e3, e3))) & (e3 = op(e3, e2))) | ~ sP113), inference(nnf_transformation, [], [f123])).
fof(f2003, plain, (~ spl140_151 | spl140_34), inference(avatar_split_clause, [], [f538, f1376, f1999])).
fof(f538, plain, ((e3 = op(e3, e3)) | ~ sP114), inference(cnf_transformation, [], [f176])).
fof(f176, plain, ((((e3 = unit) | ~ (e3 = op(e3, e3))) & (e3 = op(e3, e3))) | ~ sP114), inference(nnf_transformation, [], [f124])).
fof(f1996, plain, (~ spl140_150 | spl140_29), inference(avatar_split_clause, [], [f536, f1355, f1992])).
fof(f536, plain, ((e3 = op(e3, e4)) | ~ sP115), inference(cnf_transformation, [], [f175])).
fof(f175, plain, ((((e3 = unit) | ~ (e4 = op(e3, e3))) & (e3 = op(e3, e4))) | ~ sP115), inference(nnf_transformation, [], [f125])).
fof(f1989, plain, (~ spl140_149 | spl140_24), inference(avatar_split_clause, [], [f1053, f1334, f1985])).
fof(f1053, plain, ((e3 = op(e4, unit)) | ~ sP116), inference(definition_unfolding, [], [f534, f377])).
fof(f534, plain, ((e3 = op(e4, e0)) | ~ sP116), inference(cnf_transformation, [], [f174])).
fof(f174, plain, ((((e3 = unit) | ~ (e0 = op(e4, e3))) & (e3 = op(e4, e0))) | ~ sP116), inference(nnf_transformation, [], [f126])).
fof(f1982, plain, (~ spl140_148 | spl140_19), inference(avatar_split_clause, [], [f532, f1313, f1978])).
fof(f532, plain, ((e3 = op(e4, e1)) | ~ sP117), inference(cnf_transformation, [], [f173])).
fof(f173, plain, ((((e3 = unit) | ~ (e1 = op(e4, e3))) & (e3 = op(e4, e1))) | ~ sP117), inference(nnf_transformation, [], [f127])).
fof(f1981, plain, (~ spl140_148 | ~ spl140_7), inference(avatar_split_clause, [], [f1976, f1263, f1978])).
fof(f1976, plain, (~ (e1 = op(e4, e3)) | ~ sP117), inference(subsumption_resolution, [], [f533, f1031])).
fof(f533, plain, ((e3 = unit) | ~ (e1 = op(e4, e3)) | ~ sP117), inference(cnf_transformation, [], [f173])).
fof(f1975, plain, (~ spl140_147 | spl140_125), inference(avatar_split_clause, [], [f1051, f1758, f1971])).
fof(f1051, plain, ((e4 = op(unit, unit)) | ~ sP118), inference(definition_unfolding, [], [f530, f377, f377])).
fof(f530, plain, ((op(e0, e0) = e4) | ~ sP118), inference(cnf_transformation, [], [f172])).
fof(f172, plain, ((((e4 = unit) | ~ (e0 = op(e0, e4))) & (op(e0, e0) = e4)) | ~ sP118), inference(nnf_transformation, [], [f128])).
fof(f1968, plain, (~ spl140_146 | spl140_120), inference(avatar_split_clause, [], [f1049, f1737, f1964])).
fof(f1049, plain, ((e4 = op(unit, e1)) | ~ sP119), inference(definition_unfolding, [], [f528, f377])).
fof(f528, plain, ((e4 = op(e0, e1)) | ~ sP119), inference(cnf_transformation, [], [f171])).
fof(f171, plain, ((((e4 = unit) | ~ (e1 = op(e0, e4))) & (e4 = op(e0, e1))) | ~ sP119), inference(nnf_transformation, [], [f129])).
fof(f1961, plain, (~ spl140_145 | spl140_115), inference(avatar_split_clause, [], [f1047, f1716, f1957])).
fof(f1047, plain, ((e4 = op(unit, e2)) | ~ sP120), inference(definition_unfolding, [], [f526, f377])).
fof(f526, plain, ((e4 = op(e0, e2)) | ~ sP120), inference(cnf_transformation, [], [f170])).
fof(f170, plain, ((((e4 = unit) | ~ (e2 = op(e0, e4))) & (e4 = op(e0, e2))) | ~ sP120), inference(nnf_transformation, [], [f130])).
fof(f1954, plain, (~ spl140_144 | spl140_110), inference(avatar_split_clause, [], [f1045, f1695, f1950])).
fof(f1045, plain, ((e4 = op(unit, e3)) | ~ sP121), inference(definition_unfolding, [], [f524, f377])).
fof(f524, plain, ((e4 = op(e0, e3)) | ~ sP121), inference(cnf_transformation, [], [f169])).
fof(f169, plain, ((((e4 = unit) | ~ (e3 = op(e0, e4))) & (e4 = op(e0, e3))) | ~ sP121), inference(nnf_transformation, [], [f131])).
fof(f1946, plain, (~ spl140_143 | ~ spl140_105), inference(avatar_split_clause, [], [f1941, f1674, f1943])).
fof(f1941, plain, (~ (e4 = op(unit, e4)) | ~ sP122), inference(subsumption_resolution, [], [f1042, f1030])).
fof(f1042, plain, ((e4 = unit) | ~ (e4 = op(unit, e4)) | ~ sP122), inference(definition_unfolding, [], [f523, f377])).
fof(f523, plain, ((e4 = unit) | ~ (e4 = op(e0, e4)) | ~ sP122), inference(cnf_transformation, [], [f168])).
fof(f168, plain, ((((e4 = unit) | ~ (e4 = op(e0, e4))) & (e4 = op(e0, e4))) | ~ sP122), inference(nnf_transformation, [], [f132])).
fof(f1940, plain, (~ spl140_142 | spl140_100), inference(avatar_split_clause, [], [f1041, f1653, f1936])).
fof(f1041, plain, ((e4 = op(e1, unit)) | ~ sP123), inference(definition_unfolding, [], [f520, f377])).
fof(f520, plain, ((e4 = op(e1, e0)) | ~ sP123), inference(cnf_transformation, [], [f167])).
fof(f167, plain, ((((e4 = unit) | ~ (e0 = op(e1, e4))) & (e4 = op(e1, e0))) | ~ sP123), inference(nnf_transformation, [], [f133])).
fof(f1933, plain, (~ spl140_141 | spl140_95), inference(avatar_split_clause, [], [f518, f1632, f1929])).
fof(f518, plain, ((e4 = op(e1, e1)) | ~ sP124), inference(cnf_transformation, [], [f166])).
fof(f166, plain, ((((e4 = unit) | ~ (e1 = op(e1, e4))) & (e4 = op(e1, e1))) | ~ sP124), inference(nnf_transformation, [], [f134])).
fof(f1926, plain, (~ spl140_140 | spl140_90), inference(avatar_split_clause, [], [f516, f1611, f1922])).
fof(f516, plain, ((e4 = op(e1, e2)) | ~ sP125), inference(cnf_transformation, [], [f165])).
fof(f165, plain, ((((e4 = unit) | ~ (e2 = op(e1, e4))) & (e4 = op(e1, e2))) | ~ sP125), inference(nnf_transformation, [], [f135])).
fof(f1925, plain, (~ spl140_140 | ~ spl140_78), inference(avatar_split_clause, [], [f1920, f1561, f1922])).
fof(f1920, plain, (~ (e2 = op(e1, e4)) | ~ sP125), inference(subsumption_resolution, [], [f517, f1030])).
fof(f517, plain, ((e4 = unit) | ~ (e2 = op(e1, e4)) | ~ sP125), inference(cnf_transformation, [], [f165])).
fof(f1919, plain, (~ spl140_139 | spl140_85), inference(avatar_split_clause, [], [f514, f1590, f1915])).
fof(f514, plain, ((e4 = op(e1, e3)) | ~ sP126), inference(cnf_transformation, [], [f164])).
fof(f164, plain, ((((e4 = unit) | ~ (e3 = op(e1, e4))) & (e4 = op(e1, e3))) | ~ sP126), inference(nnf_transformation, [], [f136])).
fof(f1918, plain, (~ spl140_139 | ~ spl140_79), inference(avatar_split_clause, [], [f1913, f1565, f1915])).
fof(f1913, plain, (~ (e3 = op(e1, e4)) | ~ sP126), inference(subsumption_resolution, [], [f515, f1030])).
fof(f515, plain, ((e4 = unit) | ~ (e3 = op(e1, e4)) | ~ sP126), inference(cnf_transformation, [], [f164])).
fof(f1912, plain, (~ spl140_138 | spl140_80), inference(avatar_split_clause, [], [f512, f1569, f1908])).
fof(f512, plain, ((e4 = op(e1, e4)) | ~ sP127), inference(cnf_transformation, [], [f163])).
fof(f163, plain, ((((e4 = unit) | ~ (e4 = op(e1, e4))) & (e4 = op(e1, e4))) | ~ sP127), inference(nnf_transformation, [], [f137])).
fof(f1905, plain, (~ spl140_137 | spl140_75), inference(avatar_split_clause, [], [f1039, f1548, f1901])).
fof(f1039, plain, ((e4 = op(e2, unit)) | ~ sP128), inference(definition_unfolding, [], [f510, f377])).
fof(f510, plain, ((e4 = op(e2, e0)) | ~ sP128), inference(cnf_transformation, [], [f162])).
fof(f162, plain, ((((e4 = unit) | ~ (e0 = op(e2, e4))) & (e4 = op(e2, e0))) | ~ sP128), inference(nnf_transformation, [], [f138])).
fof(f1898, plain, (~ spl140_136 | spl140_70), inference(avatar_split_clause, [], [f508, f1527, f1894])).
fof(f508, plain, ((e4 = op(e2, e1)) | ~ sP129), inference(cnf_transformation, [], [f161])).
fof(f161, plain, ((((e4 = unit) | ~ (e1 = op(e2, e4))) & (e4 = op(e2, e1))) | ~ sP129), inference(nnf_transformation, [], [f139])).
fof(f1897, plain, (~ spl140_136 | ~ spl140_52), inference(avatar_split_clause, [], [f1892, f1452, f1894])).
fof(f1892, plain, (~ (e1 = op(e2, e4)) | ~ sP129), inference(subsumption_resolution, [], [f509, f1030])).
fof(f509, plain, ((e4 = unit) | ~ (e1 = op(e2, e4)) | ~ sP129), inference(cnf_transformation, [], [f161])).
fof(f1891, plain, (~ spl140_135 | spl140_65), inference(avatar_split_clause, [], [f506, f1506, f1887])).
fof(f506, plain, ((e4 = op(e2, e2)) | ~ sP130), inference(cnf_transformation, [], [f160])).
fof(f160, plain, ((((e4 = unit) | ~ (e2 = op(e2, e4))) & (e4 = op(e2, e2))) | ~ sP130), inference(nnf_transformation, [], [f140])).
fof(f1884, plain, (~ spl140_134 | spl140_60), inference(avatar_split_clause, [], [f504, f1485, f1880])).
fof(f504, plain, ((e4 = op(e2, e3)) | ~ sP131), inference(cnf_transformation, [], [f159])).
fof(f159, plain, ((((e4 = unit) | ~ (e3 = op(e2, e4))) & (e4 = op(e2, e3))) | ~ sP131), inference(nnf_transformation, [], [f141])).
fof(f1883, plain, (~ spl140_134 | ~ spl140_54), inference(avatar_split_clause, [], [f1878, f1460, f1880])).
fof(f1878, plain, (~ (e3 = op(e2, e4)) | ~ sP131), inference(subsumption_resolution, [], [f505, f1030])).
fof(f505, plain, ((e4 = unit) | ~ (e3 = op(e2, e4)) | ~ sP131), inference(cnf_transformation, [], [f159])).
fof(f1877, plain, (~ spl140_133 | spl140_55), inference(avatar_split_clause, [], [f502, f1464, f1873])).
fof(f502, plain, ((e4 = op(e2, e4)) | ~ sP132), inference(cnf_transformation, [], [f158])).
fof(f158, plain, ((((e4 = unit) | ~ (e4 = op(e2, e4))) & (e4 = op(e2, e4))) | ~ sP132), inference(nnf_transformation, [], [f142])).
fof(f1870, plain, (~ spl140_132 | spl140_50), inference(avatar_split_clause, [], [f1037, f1443, f1866])).
fof(f1037, plain, ((e4 = op(e3, unit)) | ~ sP133), inference(definition_unfolding, [], [f500, f377])).
fof(f500, plain, ((e4 = op(e3, e0)) | ~ sP133), inference(cnf_transformation, [], [f157])).
fof(f157, plain, ((((e4 = unit) | ~ (e0 = op(e3, e4))) & (e4 = op(e3, e0))) | ~ sP133), inference(nnf_transformation, [], [f143])).
fof(f1863, plain, (~ spl140_131 | spl140_45), inference(avatar_split_clause, [], [f498, f1422, f1859])).
fof(f498, plain, ((e4 = op(e3, e1)) | ~ sP134), inference(cnf_transformation, [], [f156])).
fof(f156, plain, ((((e4 = unit) | ~ (e1 = op(e3, e4))) & (e4 = op(e3, e1))) | ~ sP134), inference(nnf_transformation, [], [f144])).
fof(f1862, plain, (~ spl140_131 | ~ spl140_27), inference(avatar_split_clause, [], [f1857, f1347, f1859])).
fof(f1857, plain, (~ (e1 = op(e3, e4)) | ~ sP134), inference(subsumption_resolution, [], [f499, f1030])).
fof(f499, plain, ((e4 = unit) | ~ (e1 = op(e3, e4)) | ~ sP134), inference(cnf_transformation, [], [f156])).
fof(f1856, plain, (~ spl140_130 | spl140_40), inference(avatar_split_clause, [], [f496, f1401, f1852])).
fof(f496, plain, ((e4 = op(e3, e2)) | ~ sP135), inference(cnf_transformation, [], [f155])).
fof(f155, plain, ((((e4 = unit) | ~ (e2 = op(e3, e4))) & (e4 = op(e3, e2))) | ~ sP135), inference(nnf_transformation, [], [f145])).
fof(f1855, plain, (~ spl140_130 | ~ spl140_28), inference(avatar_split_clause, [], [f1850, f1351, f1852])).
fof(f1850, plain, (~ (e2 = op(e3, e4)) | ~ sP135), inference(subsumption_resolution, [], [f497, f1030])).
fof(f497, plain, ((e4 = unit) | ~ (e2 = op(e3, e4)) | ~ sP135), inference(cnf_transformation, [], [f155])).
fof(f1849, plain, (~ spl140_129 | spl140_35), inference(avatar_split_clause, [], [f494, f1380, f1845])).
fof(f494, plain, ((e4 = op(e3, e3)) | ~ sP136), inference(cnf_transformation, [], [f154])).
fof(f154, plain, ((((e4 = unit) | ~ (e3 = op(e3, e4))) & (e4 = op(e3, e3))) | ~ sP136), inference(nnf_transformation, [], [f146])).
fof(f1842, plain, (~ spl140_128 | spl140_30), inference(avatar_split_clause, [], [f492, f1359, f1838])).
fof(f492, plain, ((e4 = op(e3, e4)) | ~ sP137), inference(cnf_transformation, [], [f153])).
fof(f153, plain, ((((e4 = unit) | ~ (e4 = op(e3, e4))) & (e4 = op(e3, e4))) | ~ sP137), inference(nnf_transformation, [], [f147])).
fof(f1834, plain, (~ spl140_127 | ~ spl140_1), inference(avatar_split_clause, [], [f1829, f1238, f1831])).
fof(f1829, plain, (~ (op(e4, e4) = unit) | ~ sP138), inference(subsumption_resolution, [], [f1034, f1030])).
fof(f1034, plain, ((e4 = unit) | ~ (op(e4, e4) = unit) | ~ sP138), inference(definition_unfolding, [], [f491, f377])).
fof(f491, plain, ((e4 = unit) | ~ (e0 = op(e4, e4)) | ~ sP138), inference(cnf_transformation, [], [f152])).
fof(f152, plain, ((((e4 = unit) | ~ (e0 = op(e4, e4))) & (e4 = op(e4, e0))) | ~ sP138), inference(nnf_transformation, [], [f148])).
fof(f1828, plain, (~ spl140_126 | spl140_20), inference(avatar_split_clause, [], [f488, f1317, f1824])).
fof(f488, plain, ((e4 = op(e4, e1)) | ~ sP139), inference(cnf_transformation, [], [f151])).
fof(f151, plain, ((((e4 = unit) | ~ (e1 = op(e4, e4))) & (e4 = op(e4, e1))) | ~ sP139), inference(nnf_transformation, [], [f149])).
fof(f1813, plain, (spl140_125 | spl140_120 | spl140_115 | spl140_110 | spl140_105), inference(avatar_split_clause, [], [f969, f1674, f1695, f1716, f1737, f1758])).
fof(f969, plain, ((e4 = op(unit, e4)) | (e4 = op(unit, e3)) | (e4 = op(unit, e2)) | (e4 = op(unit, e1)) | (e4 = op(unit, unit))), inference(definition_unfolding, [], [f335, f377, f377, f377, f377, f377, f377])).
fof(f335, plain, ((e4 = op(e0, e4)) | (e4 = op(e0, e3)) | (e4 = op(e0, e2)) | (e4 = op(e0, e1)) | (op(e0, e0) = e4)), inference(cnf_transformation, [], [f3])).
fof(f3, plain, (((e4 = op(e4, e4)) | (e4 = op(e3, e4)) | (e4 = op(e2, e4)) | (e4 = op(e1, e4)) | (e4 = op(e0, e4))) & ((e4 = op(e4, e4)) | (e4 = op(e4, e3)) | (e4 = op(e4, e2)) | (e4 = op(e4, e1)) | (e4 = op(e4, e0))) & ((e3 = op(e4, e4)) | (e3 = op(e3, e4)) | (e3 = op(e2, e4)) | (e3 = op(e1, e4)) | (e3 = op(e0, e4))) & ((e3 = op(e4, e4)) | (e3 = op(e4, e3)) | (e3 = op(e4, e2)) | (e3 = op(e4, e1)) | (e3 = op(e4, e0))) & ((e2 = op(e4, e4)) | (e2 = op(e3, e4)) | (e2 = op(e2, e4)) | (e2 = op(e1, e4)) | (e2 = op(e0, e4))) & ((e2 = op(e4, e4)) | (e2 = op(e4, e3)) | (e2 = op(e4, e2)) | (e2 = op(e4, e1)) | (e2 = op(e4, e0))) & ((e1 = op(e4, e4)) | (e1 = op(e3, e4)) | (e1 = op(e2, e4)) | (e1 = op(e1, e4)) | (e1 = op(e0, e4))) & ((e1 = op(e4, e4)) | (e1 = op(e4, e3)) | (e1 = op(e4, e2)) | (e1 = op(e4, e1)) | (e1 = op(e4, e0))) & ((e0 = op(e4, e4)) | (e0 = op(e3, e4)) | (e0 = op(e2, e4)) | (e0 = op(e1, e4)) | (e0 = op(e0, e4))) & ((e0 = op(e4, e4)) | (e0 = op(e4, e3)) | (e0 = op(e4, e2)) | (e0 = op(e4, e1)) | (e0 = op(e4, e0))) & ((e4 = op(e4, e3)) | (e4 = op(e3, e3)) | (e4 = op(e2, e3)) | (e4 = op(e1, e3)) | (e4 = op(e0, e3))) & ((e4 = op(e3, e4)) | (e4 = op(e3, e3)) | (e4 = op(e3, e2)) | (e4 = op(e3, e1)) | (e4 = op(e3, e0))) & ((e3 = op(e4, e3)) | (e3 = op(e3, e3)) | (e3 = op(e2, e3)) | (e3 = op(e1, e3)) | (e3 = op(e0, e3))) & ((e3 = op(e3, e4)) | (e3 = op(e3, e3)) | (e3 = op(e3, e2)) | (e3 = op(e3, e1)) | (e3 = op(e3, e0))) & ((e2 = op(e4, e3)) | (e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(e0, e3))) & ((e2 = op(e3, e4)) | (e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, e0))) & ((e1 = op(e4, e3)) | (e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(e0, e3))) & ((e1 = op(e3, e4)) | (e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, e0))) & ((e0 = op(e4, e3)) | (e0 = op(e3, e3)) | (e0 = op(e2, e3)) | (e0 = op(e1, e3)) | (e0 = op(e0, e3))) & ((e0 = op(e3, e4)) | (e0 = op(e3, e3)) | (e0 = op(e3, e2)) | (e0 = op(e3, e1)) | (e0 = op(e3, e0))) & ((e4 = op(e4, e2)) | (e4 = op(e3, e2)) | (e4 = op(e2, e2)) | (e4 = op(e1, e2)) | (e4 = op(e0, e2))) & ((e4 = op(e2, e4)) | (e4 = op(e2, e3)) | (e4 = op(e2, e2)) | (e4 = op(e2, e1)) | (e4 = op(e2, e0))) & ((e3 = op(e4, e2)) | (e3 = op(e3, e2)) | (e3 = op(e2, e2)) | (e3 = op(e1, e2)) | (e3 = op(e0, e2))) & ((e3 = op(e2, e4)) | (e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, e0))) & ((e2 = op(e4, e2)) | (e2 = op(e3, e2)) | (e2 = op(e2, e2)) | (e2 = op(e1, e2)) | (e2 = op(e0, e2))) & ((e2 = op(e2, e4)) | (e2 = op(e2, e3)) | (e2 = op(e2, e2)) | (e2 = op(e2, e1)) | (e2 = op(e2, e0))) & ((e1 = op(e4, e2)) | (e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(e0, e2))) & ((e1 = op(e2, e4)) | (e1 = op(e2, e3)) | (e1 = op(e2, e2)) | (e1 = op(e2, e1)) | (e1 = op(e2, e0))) & ((e0 = op(e4, e2)) | (e0 = op(e3, e2)) | (e0 = op(e2, e2)) | (e0 = op(e1, e2)) | (e0 = op(e0, e2))) & ((e0 = op(e2, e4)) | (e0 = op(e2, e3)) | (e0 = op(e2, e2)) | (e0 = op(e2, e1)) | (e0 = op(e2, e0))) & ((e4 = op(e4, e1)) | (e4 = op(e3, e1)) | (e4 = op(e2, e1)) | (e4 = op(e1, e1)) | (e4 = op(e0, e1))) & ((e4 = op(e1, e4)) | (e4 = op(e1, e3)) | (e4 = op(e1, e2)) | (e4 = op(e1, e1)) | (e4 = op(e1, e0))) & ((e3 = op(e4, e1)) | (e3 = op(e3, e1)) | (e3 = op(e2, e1)) | (e3 = op(e1, e1)) | (e3 = op(e0, e1))) & ((e3 = op(e1, e4)) | (e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, e0))) & ((e2 = op(e4, e1)) | (e2 = op(e3, e1)) | (e2 = op(e2, e1)) | (e2 = op(e1, e1)) | (e2 = op(e0, e1))) & ((e2 = op(e1, e4)) | (e2 = op(e1, e3)) | (e2 = op(e1, e2)) | (e2 = op(e1, e1)) | (e2 = op(e1, e0))) & ((e1 = op(e4, e1)) | (e1 = op(e3, e1)) | (e1 = op(e2, e1)) | (e1 = op(e1, e1)) | (e1 = op(e0, e1))) & ((e1 = op(e1, e4)) | (e1 = op(e1, e3)) | (e1 = op(e1, e2)) | (e1 = op(e1, e1)) | (e1 = op(e1, e0))) & ((e0 = op(e4, e1)) | (e0 = op(e3, e1)) | (e0 = op(e2, e1)) | (e0 = op(e1, e1)) | (e0 = op(e0, e1))) & ((e0 = op(e1, e4)) | (e0 = op(e1, e3)) | (e0 = op(e1, e2)) | (e0 = op(e1, e1)) | (e0 = op(e1, e0))) & ((e4 = op(e4, e0)) | (e4 = op(e3, e0)) | (e4 = op(e2, e0)) | (e4 = op(e1, e0)) | (op(e0, e0) = e4)) & ((e4 = op(e0, e4)) | (e4 = op(e0, e3)) | (e4 = op(e0, e2)) | (e4 = op(e0, e1)) | (op(e0, e0) = e4)) & ((e3 = op(e4, e0)) | (e3 = op(e3, e0)) | (e3 = op(e2, e0)) | (e3 = op(e1, e0)) | (op(e0, e0) = e3)) & ((e3 = op(e0, e4)) | (e3 = op(e0, e3)) | (e3 = op(e0, e2)) | (e3 = op(e0, e1)) | (op(e0, e0) = e3)) & ((e2 = op(e4, e0)) | (e2 = op(e3, e0)) | (e2 = op(e2, e0)) | (e2 = op(e1, e0)) | (op(e0, e0) = e2)) & ((e2 = op(e0, e4)) | (e2 = op(e0, e3)) | (e2 = op(e0, e2)) | (e2 = op(e0, e1)) | (op(e0, e0) = e2)) & ((e1 = op(e4, e0)) | (e1 = op(e3, e0)) | (e1 = op(e2, e0)) | (e1 = op(e1, e0)) | (op(e0, e0) = e1)) & ((e1 = op(e0, e4)) | (e1 = op(e0, e3)) | (e1 = op(e0, e2)) | (e1 = op(e0, e1)) | (op(e0, e0) = e1)) & ((e0 = op(e4, e0)) | (e0 = op(e3, e0)) | (e0 = op(e2, e0)) | (e0 = op(e1, e0)) | (e0 = op(e0, e0))) & ((e0 = op(e0, e4)) | (e0 = op(e0, e3)) | (e0 = op(e0, e2)) | (e0 = op(e0, e1)) | (e0 = op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG047+1.p', ax3)).
fof(f1812, plain, (spl140_125 | spl140_100 | spl140_75 | spl140_50 | spl140_25), inference(avatar_split_clause, [], [f968, f1338, f1443, f1548, f1653, f1758])).
fof(f968, plain, ((e4 = op(e4, unit)) | (e4 = op(e3, unit)) | (e4 = op(e2, unit)) | (e4 = op(e1, unit)) | (e4 = op(unit, unit))), inference(definition_unfolding, [], [f336, f377, f377, f377, f377, f377, f377])).
fof(f336, plain, ((e4 = op(e4, e0)) | (e4 = op(e3, e0)) | (e4 = op(e2, e0)) | (e4 = op(e1, e0)) | (op(e0, e0) = e4)), inference(cnf_transformation, [], [f3])).
fof(f1811, plain, (spl140_96 | spl140_91 | spl140_86 | spl140_81 | spl140_76), inference(avatar_split_clause, [], [f967, f1553, f1574, f1595, f1616, f1637])).
fof(f967, plain, ((op(e1, e4) = unit) | (op(e1, e3) = unit) | (op(e1, e2) = unit) | (op(e1, e1) = unit) | (unit = op(e1, unit))), inference(definition_unfolding, [], [f337, f377, f377, f377, f377, f377, f377])).
fof(f337, plain, ((e0 = op(e1, e4)) | (e0 = op(e1, e3)) | (e0 = op(e1, e2)) | (e0 = op(e1, e1)) | (e0 = op(e1, e0))), inference(cnf_transformation, [], [f3])).
fof(f1810, plain, (spl140_116 | spl140_91 | spl140_66 | spl140_41 | spl140_16), inference(avatar_split_clause, [], [f966, f1301, f1406, f1511, f1616, f1721])).
fof(f966, plain, ((op(e4, e1) = unit) | (op(e3, e1) = unit) | (op(e2, e1) = unit) | (op(e1, e1) = unit) | (unit = op(unit, e1))), inference(definition_unfolding, [], [f338, f377, f377, f377, f377, f377, f377])).
fof(f338, plain, ((e0 = op(e4, e1)) | (e0 = op(e3, e1)) | (e0 = op(e2, e1)) | (e0 = op(e1, e1)) | (e0 = op(e0, e1))), inference(cnf_transformation, [], [f3])).
fof(f1807, plain, (spl140_98 | spl140_93 | spl140_88 | spl140_83 | spl140_78), inference(avatar_split_clause, [], [f963, f1561, f1582, f1603, f1624, f1645])).
fof(f963, plain, ((e2 = op(e1, e4)) | (e2 = op(e1, e3)) | (e2 = op(e1, e2)) | (e2 = op(e1, e1)) | (e2 = op(e1, unit))), inference(definition_unfolding, [], [f341, f377])).
fof(f341, plain, ((e2 = op(e1, e4)) | (e2 = op(e1, e3)) | (e2 = op(e1, e2)) | (e2 = op(e1, e1)) | (e2 = op(e1, e0))), inference(cnf_transformation, [], [f3])).
fof(f1806, plain, (spl140_118 | spl140_93 | spl140_68 | spl140_43 | spl140_18), inference(avatar_split_clause, [], [f962, f1309, f1414, f1519, f1624, f1729])).
fof(f962, plain, ((e2 = op(e4, e1)) | (e2 = op(e3, e1)) | (e2 = op(e2, e1)) | (e2 = op(e1, e1)) | (e2 = op(unit, e1))), inference(definition_unfolding, [], [f342, f377])).
fof(f342, plain, ((e2 = op(e4, e1)) | (e2 = op(e3, e1)) | (e2 = op(e2, e1)) | (e2 = op(e1, e1)) | (e2 = op(e0, e1))), inference(cnf_transformation, [], [f3])).
fof(f1805, plain, (spl140_99 | spl140_94 | spl140_89 | spl140_84 | spl140_79), inference(avatar_split_clause, [], [f961, f1565, f1586, f1607, f1628, f1649])).
fof(f961, plain, ((e3 = op(e1, e4)) | (e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, unit))), inference(definition_unfolding, [], [f343, f377])).
fof(f343, plain, ((e3 = op(e1, e4)) | (e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, e0))), inference(cnf_transformation, [], [f3])).
fof(f1804, plain, (spl140_119 | spl140_94 | spl140_69 | spl140_44 | spl140_19), inference(avatar_split_clause, [], [f960, f1313, f1418, f1523, f1628, f1733])).
fof(f960, plain, ((e3 = op(e4, e1)) | (e3 = op(e3, e1)) | (e3 = op(e2, e1)) | (e3 = op(e1, e1)) | (e3 = op(unit, e1))), inference(definition_unfolding, [], [f344, f377])).
fof(f344, plain, ((e3 = op(e4, e1)) | (e3 = op(e3, e1)) | (e3 = op(e2, e1)) | (e3 = op(e1, e1)) | (e3 = op(e0, e1))), inference(cnf_transformation, [], [f3])).
fof(f1803, plain, (spl140_100 | spl140_95 | spl140_90 | spl140_85 | spl140_80), inference(avatar_split_clause, [], [f959, f1569, f1590, f1611, f1632, f1653])).
fof(f959, plain, ((e4 = op(e1, e4)) | (e4 = op(e1, e3)) | (e4 = op(e1, e2)) | (e4 = op(e1, e1)) | (e4 = op(e1, unit))), inference(definition_unfolding, [], [f345, f377])).
fof(f345, plain, ((e4 = op(e1, e4)) | (e4 = op(e1, e3)) | (e4 = op(e1, e2)) | (e4 = op(e1, e1)) | (e4 = op(e1, e0))), inference(cnf_transformation, [], [f3])).
fof(f1802, plain, (spl140_120 | spl140_95 | spl140_70 | spl140_45 | spl140_20), inference(avatar_split_clause, [], [f958, f1317, f1422, f1527, f1632, f1737])).
fof(f958, plain, ((e4 = op(e4, e1)) | (e4 = op(e3, e1)) | (e4 = op(e2, e1)) | (e4 = op(e1, e1)) | (e4 = op(unit, e1))), inference(definition_unfolding, [], [f346, f377])).
fof(f346, plain, ((e4 = op(e4, e1)) | (e4 = op(e3, e1)) | (e4 = op(e2, e1)) | (e4 = op(e1, e1)) | (e4 = op(e0, e1))), inference(cnf_transformation, [], [f3])).
fof(f1799, plain, (spl140_72 | spl140_67 | spl140_62 | spl140_57 | spl140_52), inference(avatar_split_clause, [], [f955, f1452, f1473, f1494, f1515, f1536])).
fof(f955, plain, ((e1 = op(e2, e4)) | (e1 = op(e2, e3)) | (e1 = op(e2, e2)) | (e1 = op(e2, e1)) | (e1 = op(e2, unit))), inference(definition_unfolding, [], [f349, f377])).
fof(f349, plain, ((e1 = op(e2, e4)) | (e1 = op(e2, e3)) | (e1 = op(e2, e2)) | (e1 = op(e2, e1)) | (e1 = op(e2, e0))), inference(cnf_transformation, [], [f3])).
fof(f1798, plain, (spl140_112 | spl140_87 | spl140_62 | spl140_37 | spl140_12), inference(avatar_split_clause, [], [f954, f1284, f1389, f1494, f1599, f1704])).
fof(f954, plain, ((e1 = op(e4, e2)) | (e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(unit, e2))), inference(definition_unfolding, [], [f350, f377])).
fof(f350, plain, ((e1 = op(e4, e2)) | (e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(e0, e2))), inference(cnf_transformation, [], [f3])).
fof(f1795, plain, (spl140_74 | spl140_69 | spl140_64 | spl140_59 | spl140_54), inference(avatar_split_clause, [], [f951, f1460, f1481, f1502, f1523, f1544])).
fof(f951, plain, ((e3 = op(e2, e4)) | (e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, unit))), inference(definition_unfolding, [], [f353, f377])).
fof(f353, plain, ((e3 = op(e2, e4)) | (e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, e0))), inference(cnf_transformation, [], [f3])).
fof(f1794, plain, (spl140_114 | spl140_89 | spl140_64 | spl140_39 | spl140_14), inference(avatar_split_clause, [], [f950, f1292, f1397, f1502, f1607, f1712])).
fof(f950, plain, ((e3 = op(e4, e2)) | (e3 = op(e3, e2)) | (e3 = op(e2, e2)) | (e3 = op(e1, e2)) | (e3 = op(unit, e2))), inference(definition_unfolding, [], [f354, f377])).
fof(f354, plain, ((e3 = op(e4, e2)) | (e3 = op(e3, e2)) | (e3 = op(e2, e2)) | (e3 = op(e1, e2)) | (e3 = op(e0, e2))), inference(cnf_transformation, [], [f3])).
fof(f1793, plain, (spl140_75 | spl140_70 | spl140_65 | spl140_60 | spl140_55), inference(avatar_split_clause, [], [f949, f1464, f1485, f1506, f1527, f1548])).
fof(f949, plain, ((e4 = op(e2, e4)) | (e4 = op(e2, e3)) | (e4 = op(e2, e2)) | (e4 = op(e2, e1)) | (e4 = op(e2, unit))), inference(definition_unfolding, [], [f355, f377])).
fof(f355, plain, ((e4 = op(e2, e4)) | (e4 = op(e2, e3)) | (e4 = op(e2, e2)) | (e4 = op(e2, e1)) | (e4 = op(e2, e0))), inference(cnf_transformation, [], [f3])).
fof(f1792, plain, (spl140_115 | spl140_90 | spl140_65 | spl140_40 | spl140_15), inference(avatar_split_clause, [], [f948, f1296, f1401, f1506, f1611, f1716])).
fof(f948, plain, ((e4 = op(e4, e2)) | (e4 = op(e3, e2)) | (e4 = op(e2, e2)) | (e4 = op(e1, e2)) | (e4 = op(unit, e2))), inference(definition_unfolding, [], [f356, f377])).
fof(f356, plain, ((e4 = op(e4, e2)) | (e4 = op(e3, e2)) | (e4 = op(e2, e2)) | (e4 = op(e1, e2)) | (e4 = op(e0, e2))), inference(cnf_transformation, [], [f3])).
fof(f1791, plain, (spl140_46 | spl140_41 | spl140_36 | spl140_31 | spl140_26), inference(avatar_split_clause, [], [f947, f1343, f1364, f1385, f1406, f1427])).
fof(f947, plain, ((op(e3, e4) = unit) | (op(e3, e3) = unit) | (op(e3, e2) = unit) | (op(e3, e1) = unit) | (unit = op(e3, unit))), inference(definition_unfolding, [], [f357, f377, f377, f377, f377, f377, f377])).
fof(f357, plain, ((e0 = op(e3, e4)) | (e0 = op(e3, e3)) | (e0 = op(e3, e2)) | (e0 = op(e3, e1)) | (e0 = op(e3, e0))), inference(cnf_transformation, [], [f3])).
fof(f1790, plain, (spl140_106 | spl140_81 | spl140_56 | spl140_31 | spl140_6), inference(avatar_split_clause, [], [f946, f1259, f1364, f1469, f1574, f1679])).
fof(f946, plain, ((op(e4, e3) = unit) | (op(e3, e3) = unit) | (op(e2, e3) = unit) | (op(e1, e3) = unit) | (unit = op(unit, e3))), inference(definition_unfolding, [], [f358, f377, f377, f377, f377, f377, f377])).
fof(f358, plain, ((e0 = op(e4, e3)) | (e0 = op(e3, e3)) | (e0 = op(e2, e3)) | (e0 = op(e1, e3)) | (e0 = op(e0, e3))), inference(cnf_transformation, [], [f3])).
fof(f1789, plain, (spl140_47 | spl140_42 | spl140_37 | spl140_32 | spl140_27), inference(avatar_split_clause, [], [f945, f1347, f1368, f1389, f1410, f1431])).
fof(f945, plain, ((e1 = op(e3, e4)) | (e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, unit))), inference(definition_unfolding, [], [f359, f377])).
fof(f359, plain, ((e1 = op(e3, e4)) | (e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, e0))), inference(cnf_transformation, [], [f3])).
fof(f1788, plain, (spl140_107 | spl140_82 | spl140_57 | spl140_32 | spl140_7), inference(avatar_split_clause, [], [f944, f1263, f1368, f1473, f1578, f1683])).
fof(f944, plain, ((e1 = op(e4, e3)) | (e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(unit, e3))), inference(definition_unfolding, [], [f360, f377])).
fof(f360, plain, ((e1 = op(e4, e3)) | (e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(e0, e3))), inference(cnf_transformation, [], [f3])).
fof(f1787, plain, (spl140_48 | spl140_43 | spl140_38 | spl140_33 | spl140_28), inference(avatar_split_clause, [], [f943, f1351, f1372, f1393, f1414, f1435])).
fof(f943, plain, ((e2 = op(e3, e4)) | (e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, unit))), inference(definition_unfolding, [], [f361, f377])).
fof(f361, plain, ((e2 = op(e3, e4)) | (e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, e0))), inference(cnf_transformation, [], [f3])).
fof(f1786, plain, (spl140_108 | spl140_83 | spl140_58 | spl140_33 | spl140_8), inference(avatar_split_clause, [], [f942, f1267, f1372, f1477, f1582, f1687])).
fof(f942, plain, ((e2 = op(e4, e3)) | (e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(unit, e3))), inference(definition_unfolding, [], [f362, f377])).
fof(f362, plain, ((e2 = op(e4, e3)) | (e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(e0, e3))), inference(cnf_transformation, [], [f3])).
fof(f1783, plain, (spl140_50 | spl140_45 | spl140_40 | spl140_35 | spl140_30), inference(avatar_split_clause, [], [f939, f1359, f1380, f1401, f1422, f1443])).
fof(f939, plain, ((e4 = op(e3, e4)) | (e4 = op(e3, e3)) | (e4 = op(e3, e2)) | (e4 = op(e3, e1)) | (e4 = op(e3, unit))), inference(definition_unfolding, [], [f365, f377])).
fof(f365, plain, ((e4 = op(e3, e4)) | (e4 = op(e3, e3)) | (e4 = op(e3, e2)) | (e4 = op(e3, e1)) | (e4 = op(e3, e0))), inference(cnf_transformation, [], [f3])).
fof(f1782, plain, (spl140_110 | spl140_85 | spl140_60 | spl140_35 | spl140_10), inference(avatar_split_clause, [], [f938, f1275, f1380, f1485, f1590, f1695])).
fof(f938, plain, ((e4 = op(e4, e3)) | (e4 = op(e3, e3)) | (e4 = op(e2, e3)) | (e4 = op(e1, e3)) | (e4 = op(unit, e3))), inference(definition_unfolding, [], [f366, f377])).
fof(f366, plain, ((e4 = op(e4, e3)) | (e4 = op(e3, e3)) | (e4 = op(e2, e3)) | (e4 = op(e1, e3)) | (e4 = op(e0, e3))), inference(cnf_transformation, [], [f3])).
fof(f1781, plain, (spl140_21 | spl140_16 | spl140_11 | spl140_6 | spl140_1), inference(avatar_split_clause, [], [f937, f1238, f1259, f1280, f1301, f1322])).
fof(f937, plain, ((op(e4, e4) = unit) | (op(e4, e3) = unit) | (op(e4, e2) = unit) | (op(e4, e1) = unit) | (unit = op(e4, unit))), inference(definition_unfolding, [], [f367, f377, f377, f377, f377, f377, f377])).
fof(f367, plain, ((e0 = op(e4, e4)) | (e0 = op(e4, e3)) | (e0 = op(e4, e2)) | (e0 = op(e4, e1)) | (e0 = op(e4, e0))), inference(cnf_transformation, [], [f3])).
fof(f1780, plain, (spl140_101 | spl140_76 | spl140_51 | spl140_26 | spl140_1), inference(avatar_split_clause, [], [f936, f1238, f1343, f1448, f1553, f1658])).
fof(f936, plain, ((op(e4, e4) = unit) | (op(e3, e4) = unit) | (op(e2, e4) = unit) | (op(e1, e4) = unit) | (unit = op(unit, e4))), inference(definition_unfolding, [], [f368, f377, f377, f377, f377, f377, f377])).
fof(f368, plain, ((e0 = op(e4, e4)) | (e0 = op(e3, e4)) | (e0 = op(e2, e4)) | (e0 = op(e1, e4)) | (e0 = op(e0, e4))), inference(cnf_transformation, [], [f3])).
fof(f1779, plain, (spl140_22 | spl140_17 | spl140_12 | spl140_7 | spl140_2), inference(avatar_split_clause, [], [f935, f1242, f1263, f1284, f1305, f1326])).
fof(f935, plain, ((e1 = op(e4, e4)) | (e1 = op(e4, e3)) | (e1 = op(e4, e2)) | (e1 = op(e4, e1)) | (e1 = op(e4, unit))), inference(definition_unfolding, [], [f369, f377])).
fof(f369, plain, ((e1 = op(e4, e4)) | (e1 = op(e4, e3)) | (e1 = op(e4, e2)) | (e1 = op(e4, e1)) | (e1 = op(e4, e0))), inference(cnf_transformation, [], [f3])).
fof(f1778, plain, (spl140_102 | spl140_77 | spl140_52 | spl140_27 | spl140_2), inference(avatar_split_clause, [], [f934, f1242, f1347, f1452, f1557, f1662])).
fof(f934, plain, ((e1 = op(e4, e4)) | (e1 = op(e3, e4)) | (e1 = op(e2, e4)) | (e1 = op(e1, e4)) | (e1 = op(unit, e4))), inference(definition_unfolding, [], [f370, f377])).
fof(f370, plain, ((e1 = op(e4, e4)) | (e1 = op(e3, e4)) | (e1 = op(e2, e4)) | (e1 = op(e1, e4)) | (e1 = op(e0, e4))), inference(cnf_transformation, [], [f3])).
fof(f1777, plain, (spl140_23 | spl140_18 | spl140_13 | spl140_8 | spl140_3), inference(avatar_split_clause, [], [f933, f1246, f1267, f1288, f1309, f1330])).
fof(f933, plain, ((e2 = op(e4, e4)) | (e2 = op(e4, e3)) | (e2 = op(e4, e2)) | (e2 = op(e4, e1)) | (e2 = op(e4, unit))), inference(definition_unfolding, [], [f371, f377])).
fof(f371, plain, ((e2 = op(e4, e4)) | (e2 = op(e4, e3)) | (e2 = op(e4, e2)) | (e2 = op(e4, e1)) | (e2 = op(e4, e0))), inference(cnf_transformation, [], [f3])).
fof(f1776, plain, (spl140_103 | spl140_78 | spl140_53 | spl140_28 | spl140_3), inference(avatar_split_clause, [], [f932, f1246, f1351, f1456, f1561, f1666])).
fof(f932, plain, ((e2 = op(e4, e4)) | (e2 = op(e3, e4)) | (e2 = op(e2, e4)) | (e2 = op(e1, e4)) | (e2 = op(unit, e4))), inference(definition_unfolding, [], [f372, f377])).
fof(f372, plain, ((e2 = op(e4, e4)) | (e2 = op(e3, e4)) | (e2 = op(e2, e4)) | (e2 = op(e1, e4)) | (e2 = op(e0, e4))), inference(cnf_transformation, [], [f3])).
fof(f1775, plain, (spl140_24 | spl140_19 | spl140_14 | spl140_9 | spl140_4), inference(avatar_split_clause, [], [f931, f1250, f1271, f1292, f1313, f1334])).
fof(f931, plain, ((e3 = op(e4, e4)) | (e3 = op(e4, e3)) | (e3 = op(e4, e2)) | (e3 = op(e4, e1)) | (e3 = op(e4, unit))), inference(definition_unfolding, [], [f373, f377])).
fof(f373, plain, ((e3 = op(e4, e4)) | (e3 = op(e4, e3)) | (e3 = op(e4, e2)) | (e3 = op(e4, e1)) | (e3 = op(e4, e0))), inference(cnf_transformation, [], [f3])).
fof(f1774, plain, (spl140_104 | spl140_79 | spl140_54 | spl140_29 | spl140_4), inference(avatar_split_clause, [], [f930, f1250, f1355, f1460, f1565, f1670])).
fof(f930, plain, ((e3 = op(e4, e4)) | (e3 = op(e3, e4)) | (e3 = op(e2, e4)) | (e3 = op(e1, e4)) | (e3 = op(unit, e4))), inference(definition_unfolding, [], [f374, f377])).
fof(f374, plain, ((e3 = op(e4, e4)) | (e3 = op(e3, e4)) | (e3 = op(e2, e4)) | (e3 = op(e1, e4)) | (e3 = op(e0, e4))), inference(cnf_transformation, [], [f3])).
fof(f1771, plain, spl140_121, inference(avatar_split_clause, [], [f927, f1742])).
fof(f927, plain, (unit = op(unit, unit)), inference(definition_unfolding, [], [f316, f377, f377])).
fof(f316, plain, (e0 = op(unit, e0)), inference(cnf_transformation, [], [f2])).
fof(f2, plain, (((e4 = unit) | (e3 = unit) | (e2 = unit) | (e1 = unit) | (e0 = unit)) & (e4 = op(e4, unit)) & (e4 = op(unit, e4)) & (e3 = op(e3, unit)) & (e3 = op(unit, e3)) & (e2 = op(e2, unit)) & (e2 = op(unit, e2)) & (e1 = op(e1, unit)) & (e1 = op(unit, e1)) & (e0 = op(e0, unit)) & (e0 = op(unit, e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG047+1.p', ax2)).
fof(f1769, plain, spl140_117, inference(avatar_split_clause, [], [f318, f1725])).
fof(f318, plain, (e1 = op(unit, e1)), inference(cnf_transformation, [], [f2])).
fof(f1768, plain, spl140_97, inference(avatar_split_clause, [], [f319, f1641])).
fof(f319, plain, (e1 = op(e1, unit)), inference(cnf_transformation, [], [f2])).
fof(f1767, plain, spl140_113, inference(avatar_split_clause, [], [f320, f1708])).
fof(f320, plain, (e2 = op(unit, e2)), inference(cnf_transformation, [], [f2])).
fof(f1766, plain, spl140_73, inference(avatar_split_clause, [], [f321, f1540])).
fof(f321, plain, (e2 = op(e2, unit)), inference(cnf_transformation, [], [f2])).
fof(f1765, plain, spl140_109, inference(avatar_split_clause, [], [f322, f1691])).
fof(f322, plain, (e3 = op(unit, e3)), inference(cnf_transformation, [], [f2])).
fof(f1764, plain, spl140_49, inference(avatar_split_clause, [], [f323, f1439])).
fof(f323, plain, (e3 = op(e3, unit)), inference(cnf_transformation, [], [f2])).
fof(f1635, plain, (spl140_91 | spl140_92 | spl140_93 | spl140_94 | spl140_95), inference(avatar_split_clause, [], [f918, f1632, f1628, f1624, f1620, f1616])).
fof(f918, plain, ((e4 = op(e1, e1)) | (e3 = op(e1, e1)) | (e2 = op(e1, e1)) | (e1 = op(e1, e1)) | (op(e1, e1) = unit)), inference(definition_unfolding, [], [f297, f377])).
fof(f297, plain, ((e4 = op(e1, e1)) | (e3 = op(e1, e1)) | (e2 = op(e1, e1)) | (e1 = op(e1, e1)) | (e0 = op(e1, e1))), inference(cnf_transformation, [], [f1])).
fof(f1, plain, (((e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | (e1 = op(e4, e4)) | (e0 = op(e4, e4))) & ((e4 = op(e4, e3)) | (e3 = op(e4, e3)) | (e2 = op(e4, e3)) | (e1 = op(e4, e3)) | (e0 = op(e4, e3))) & ((e4 = op(e4, e2)) | (e3 = op(e4, e2)) | (e2 = op(e4, e2)) | (e1 = op(e4, e2)) | (e0 = op(e4, e2))) & ((e4 = op(e4, e1)) | (e3 = op(e4, e1)) | (e2 = op(e4, e1)) | (e1 = op(e4, e1)) | (e0 = op(e4, e1))) & ((e4 = op(e4, e0)) | (e3 = op(e4, e0)) | (e2 = op(e4, e0)) | (e1 = op(e4, e0)) | (e0 = op(e4, e0))) & ((e4 = op(e3, e4)) | (e3 = op(e3, e4)) | (e2 = op(e3, e4)) | (e1 = op(e3, e4)) | (e0 = op(e3, e4))) & ((e4 = op(e3, e3)) | (e3 = op(e3, e3)) | (e2 = op(e3, e3)) | (e1 = op(e3, e3)) | (e0 = op(e3, e3))) & ((e4 = op(e3, e2)) | (e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (e0 = op(e3, e2))) & ((e4 = op(e3, e1)) | (e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (e0 = op(e3, e1))) & ((e4 = op(e3, e0)) | (e3 = op(e3, e0)) | (e2 = op(e3, e0)) | (e1 = op(e3, e0)) | (e0 = op(e3, e0))) & ((e4 = op(e2, e4)) | (e3 = op(e2, e4)) | (e2 = op(e2, e4)) | (e1 = op(e2, e4)) | (e0 = op(e2, e4))) & ((e4 = op(e2, e3)) | (e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))) & ((e4 = op(e2, e2)) | (e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (e0 = op(e2, e2))) & ((e4 = op(e2, e1)) | (e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))) & ((e4 = op(e2, e0)) | (e3 = op(e2, e0)) | (e2 = op(e2, e0)) | (e1 = op(e2, e0)) | (e0 = op(e2, e0))) & ((e4 = op(e1, e4)) | (e3 = op(e1, e4)) | (e2 = op(e1, e4)) | (e1 = op(e1, e4)) | (e0 = op(e1, e4))) & ((e4 = op(e1, e3)) | (e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (e0 = op(e1, e3))) & ((e4 = op(e1, e2)) | (e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (e0 = op(e1, e2))) & ((e4 = op(e1, e1)) | (e3 = op(e1, e1)) | (e2 = op(e1, e1)) | (e1 = op(e1, e1)) | (e0 = op(e1, e1))) & ((e4 = op(e1, e0)) | (e3 = op(e1, e0)) | (e2 = op(e1, e0)) | (e1 = op(e1, e0)) | (e0 = op(e1, e0))) & ((e4 = op(e0, e4)) | (e3 = op(e0, e4)) | (e2 = op(e0, e4)) | (e1 = op(e0, e4)) | (e0 = op(e0, e4))) & ((e4 = op(e0, e3)) | (e3 = op(e0, e3)) | (e2 = op(e0, e3)) | (e1 = op(e0, e3)) | (e0 = op(e0, e3))) & ((e4 = op(e0, e2)) | (e3 = op(e0, e2)) | (e2 = op(e0, e2)) | (e1 = op(e0, e2)) | (e0 = op(e0, e2))) & ((e4 = op(e0, e1)) | (e3 = op(e0, e1)) | (e2 = op(e0, e1)) | (e1 = op(e0, e1)) | (e0 = op(e0, e1))) & ((op(e0, e0) = e4) | (op(e0, e0) = e3) | (op(e0, e0) = e2) | (op(e0, e0) = e1) | (e0 = op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG047+1.p', ax1)).
fof(f1614, plain, (spl140_86 | spl140_87 | spl140_88 | spl140_89 | spl140_90), inference(avatar_split_clause, [], [f917, f1611, f1607, f1603, f1599, f1595])).
fof(f917, plain, ((e4 = op(e1, e2)) | (e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (op(e1, e2) = unit)), inference(definition_unfolding, [], [f298, f377])).
fof(f298, plain, ((e4 = op(e1, e2)) | (e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (e0 = op(e1, e2))), inference(cnf_transformation, [], [f1])).
fof(f1593, plain, (spl140_81 | spl140_82 | spl140_83 | spl140_84 | spl140_85), inference(avatar_split_clause, [], [f916, f1590, f1586, f1582, f1578, f1574])).
fof(f916, plain, ((e4 = op(e1, e3)) | (e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (op(e1, e3) = unit)), inference(definition_unfolding, [], [f299, f377])).
fof(f299, plain, ((e4 = op(e1, e3)) | (e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (e0 = op(e1, e3))), inference(cnf_transformation, [], [f1])).
fof(f1572, plain, (spl140_76 | spl140_77 | spl140_78 | spl140_79 | spl140_80), inference(avatar_split_clause, [], [f915, f1569, f1565, f1561, f1557, f1553])).
fof(f915, plain, ((e4 = op(e1, e4)) | (e3 = op(e1, e4)) | (e2 = op(e1, e4)) | (e1 = op(e1, e4)) | (op(e1, e4) = unit)), inference(definition_unfolding, [], [f300, f377])).
fof(f300, plain, ((e4 = op(e1, e4)) | (e3 = op(e1, e4)) | (e2 = op(e1, e4)) | (e1 = op(e1, e4)) | (e0 = op(e1, e4))), inference(cnf_transformation, [], [f1])).
fof(f1530, plain, (spl140_66 | spl140_67 | spl140_68 | spl140_69 | spl140_70), inference(avatar_split_clause, [], [f913, f1527, f1523, f1519, f1515, f1511])).
fof(f913, plain, ((e4 = op(e2, e1)) | (e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (op(e2, e1) = unit)), inference(definition_unfolding, [], [f302, f377])).
fof(f302, plain, ((e4 = op(e2, e1)) | (e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))), inference(cnf_transformation, [], [f1])).
fof(f1509, plain, (spl140_61 | spl140_62 | spl140_63 | spl140_64 | spl140_65), inference(avatar_split_clause, [], [f912, f1506, f1502, f1498, f1494, f1490])).
fof(f912, plain, ((e4 = op(e2, e2)) | (e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (op(e2, e2) = unit)), inference(definition_unfolding, [], [f303, f377])).
fof(f303, plain, ((e4 = op(e2, e2)) | (e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (e0 = op(e2, e2))), inference(cnf_transformation, [], [f1])).
fof(f1488, plain, (spl140_56 | spl140_57 | spl140_58 | spl140_59 | spl140_60), inference(avatar_split_clause, [], [f911, f1485, f1481, f1477, f1473, f1469])).
fof(f911, plain, ((e4 = op(e2, e3)) | (e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (op(e2, e3) = unit)), inference(definition_unfolding, [], [f304, f377])).
fof(f304, plain, ((e4 = op(e2, e3)) | (e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))), inference(cnf_transformation, [], [f1])).
fof(f1467, plain, (spl140_51 | spl140_52 | spl140_53 | spl140_54 | spl140_55), inference(avatar_split_clause, [], [f910, f1464, f1460, f1456, f1452, f1448])).
fof(f910, plain, ((e4 = op(e2, e4)) | (e3 = op(e2, e4)) | (e2 = op(e2, e4)) | (e1 = op(e2, e4)) | (op(e2, e4) = unit)), inference(definition_unfolding, [], [f305, f377])).
fof(f305, plain, ((e4 = op(e2, e4)) | (e3 = op(e2, e4)) | (e2 = op(e2, e4)) | (e1 = op(e2, e4)) | (e0 = op(e2, e4))), inference(cnf_transformation, [], [f1])).
fof(f1425, plain, (spl140_41 | spl140_42 | spl140_43 | spl140_44 | spl140_45), inference(avatar_split_clause, [], [f908, f1422, f1418, f1414, f1410, f1406])).
fof(f908, plain, ((e4 = op(e3, e1)) | (e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (op(e3, e1) = unit)), inference(definition_unfolding, [], [f307, f377])).
fof(f307, plain, ((e4 = op(e3, e1)) | (e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (e0 = op(e3, e1))), inference(cnf_transformation, [], [f1])).
fof(f1404, plain, (spl140_36 | spl140_37 | spl140_38 | spl140_39 | spl140_40), inference(avatar_split_clause, [], [f907, f1401, f1397, f1393, f1389, f1385])).
fof(f907, plain, ((e4 = op(e3, e2)) | (e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (op(e3, e2) = unit)), inference(definition_unfolding, [], [f308, f377])).
fof(f308, plain, ((e4 = op(e3, e2)) | (e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (e0 = op(e3, e2))), inference(cnf_transformation, [], [f1])).
fof(f1383, plain, (spl140_31 | spl140_32 | spl140_33 | spl140_34 | spl140_35), inference(avatar_split_clause, [], [f906, f1380, f1376, f1372, f1368, f1364])).
fof(f906, plain, ((e4 = op(e3, e3)) | (e3 = op(e3, e3)) | (e2 = op(e3, e3)) | (e1 = op(e3, e3)) | (op(e3, e3) = unit)), inference(definition_unfolding, [], [f309, f377])).
fof(f309, plain, ((e4 = op(e3, e3)) | (e3 = op(e3, e3)) | (e2 = op(e3, e3)) | (e1 = op(e3, e3)) | (e0 = op(e3, e3))), inference(cnf_transformation, [], [f1])).
fof(f1362, plain, (spl140_26 | spl140_27 | spl140_28 | spl140_29 | spl140_30), inference(avatar_split_clause, [], [f905, f1359, f1355, f1351, f1347, f1343])).
fof(f905, plain, ((e4 = op(e3, e4)) | (e3 = op(e3, e4)) | (e2 = op(e3, e4)) | (e1 = op(e3, e4)) | (op(e3, e4) = unit)), inference(definition_unfolding, [], [f310, f377])).
fof(f310, plain, ((e4 = op(e3, e4)) | (e3 = op(e3, e4)) | (e2 = op(e3, e4)) | (e1 = op(e3, e4)) | (e0 = op(e3, e4))), inference(cnf_transformation, [], [f1])).
fof(f1320, plain, (spl140_16 | spl140_17 | spl140_18 | spl140_19 | spl140_20), inference(avatar_split_clause, [], [f903, f1317, f1313, f1309, f1305, f1301])).
fof(f903, plain, ((e4 = op(e4, e1)) | (e3 = op(e4, e1)) | (e2 = op(e4, e1)) | (e1 = op(e4, e1)) | (op(e4, e1) = unit)), inference(definition_unfolding, [], [f312, f377])).
fof(f312, plain, ((e4 = op(e4, e1)) | (e3 = op(e4, e1)) | (e2 = op(e4, e1)) | (e1 = op(e4, e1)) | (e0 = op(e4, e1))), inference(cnf_transformation, [], [f1])).
fof(f1299, plain, (spl140_11 | spl140_12 | spl140_13 | spl140_14 | spl140_15), inference(avatar_split_clause, [], [f902, f1296, f1292, f1288, f1284, f1280])).
fof(f902, plain, ((e4 = op(e4, e2)) | (e3 = op(e4, e2)) | (e2 = op(e4, e2)) | (e1 = op(e4, e2)) | (op(e4, e2) = unit)), inference(definition_unfolding, [], [f313, f377])).
fof(f313, plain, ((e4 = op(e4, e2)) | (e3 = op(e4, e2)) | (e2 = op(e4, e2)) | (e1 = op(e4, e2)) | (e0 = op(e4, e2))), inference(cnf_transformation, [], [f1])).
fof(f1278, plain, (spl140_6 | spl140_7 | spl140_8 | spl140_9 | spl140_10), inference(avatar_split_clause, [], [f901, f1275, f1271, f1267, f1263, f1259])).
fof(f901, plain, ((e4 = op(e4, e3)) | (e3 = op(e4, e3)) | (e2 = op(e4, e3)) | (e1 = op(e4, e3)) | (op(e4, e3) = unit)), inference(definition_unfolding, [], [f314, f377])).
fof(f314, plain, ((e4 = op(e4, e3)) | (e3 = op(e4, e3)) | (e2 = op(e4, e3)) | (e1 = op(e4, e3)) | (e0 = op(e4, e3))), inference(cnf_transformation, [], [f1])).
fof(f1257, plain, (spl140_1 | spl140_2 | spl140_3 | spl140_4 | spl140_5), inference(avatar_split_clause, [], [f900, f1254, f1250, f1246, f1242, f1238])).
fof(f900, plain, ((e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | (e1 = op(e4, e4)) | (op(e4, e4) = unit)), inference(definition_unfolding, [], [f315, f377])).
fof(f315, plain, ((e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | (e1 = op(e4, e4)) | (e0 = op(e4, e4))), inference(cnf_transformation, [], [f1])).