fof(f849, plain, $false, inference(subsumption_resolution, [], [f848, f611])).
fof(f611, plain, ~ sP0, inference(subsumption_resolution, [], [f610, f348])).
fof(f348, plain, ~ (e3 = unit), inference(definition_unfolding, [], [f115, f135])).
fof(f135, plain, (e0 = unit), inference(cnf_transformation, [], [f3])).
fof(f3, plain, (e0 = unit), file('/home/ubuntu/library/tptp/Problems/ALG/ALG044+1.p', ax3)).
fof(f115, plain, ~ (e0 = e3), inference(cnf_transformation, [], [f1])).
fof(f1, plain, (~ (e2 = e3) & ~ (e1 = e3) & ~ (e1 = e2) & ~ (e0 = e3) & ~ (e0 = e2) & ~ (e0 = e1)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG044+1.p', ax1)).
fof(f610, plain, ((e3 = unit) | ~ sP0), inference(forward_demodulation, [], [f476, f129])).
fof(f129, plain, (e3 = op(e2, e2)), inference(cnf_transformation, [], [f2])).
fof(f2, plain, ((e0 = op(e3, e3)) & (e1 = op(e3, e2)) & (e2 = op(e3, e1)) & (e3 = op(e3, e0)) & (e1 = op(e2, e3)) & (e3 = op(e2, e2)) & (e0 = op(e2, e1)) & (e2 = op(e2, e0)) & (e2 = op(e1, e3)) & (e0 = op(e1, e2)) & (e3 = op(e1, e1)) & (e1 = op(e1, e0)) & (e3 = op(e0, e3)) & (e2 = op(e0, e2)) & (e1 = op(e0, e1)) & (e0 = op(e0, e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG044+1.p', ax2)).
fof(f476, plain, ((op(e2, e2) = unit) | ~ sP0), inference(definition_unfolding, [], [f342, f135])).
fof(f342, plain, ((e0 = op(e2, e2)) | ~ sP0), inference(cnf_transformation, [], [f112])).
fof(f112, plain, (((e0 = op(e3, e3)) & (e0 = op(e2, e2)) & (e0 = op(e1, e1)) & (e0 = op(e0, e0))) | ~ sP0), inference(nnf_transformation, [], [f8])).
fof(f8, plain, (((e0 = op(e3, e3)) & (e0 = op(e2, e2)) & (e0 = op(e1, e1)) & (e0 = op(e0, e0))) | ~ sP0), inference(usedef, [], [e8])).
fof(e8, plain, (sP0 <=> ((e0 = op(e3, e3)) & (e0 = op(e2, e2)) & (e0 = op(e1, e1)) & (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f848, plain, sP0, inference(subsumption_resolution, [], [f847, f609])).
fof(f609, plain, ~ sP1, inference(subsumption_resolution, [], [f608, f350])).
fof(f350, plain, ~ (e1 = unit), inference(definition_unfolding, [], [f113, f135])).
fof(f113, plain, ~ (e0 = e1), inference(cnf_transformation, [], [f1])).
fof(f608, plain, ((e1 = unit) | ~ sP1), inference(forward_demodulation, [], [f339, f351])).
fof(f351, plain, (op(e3, e3) = unit), inference(definition_unfolding, [], [f134, f135])).
fof(f134, plain, (e0 = op(e3, e3)), inference(cnf_transformation, [], [f2])).
fof(f339, plain, ((e1 = op(e3, e3)) | ~ sP1), inference(cnf_transformation, [], [f111])).
fof(f111, plain, (((e1 = op(e3, e3)) & (e1 = op(e2, e2)) & (e1 = op(e1, e1)) & (e1 = op(e0, e0))) | ~ sP1), inference(nnf_transformation, [], [f9])).
fof(f9, plain, (((e1 = op(e3, e3)) & (e1 = op(e2, e2)) & (e1 = op(e1, e1)) & (e1 = op(e0, e0))) | ~ sP1), inference(usedef, [], [e9])).
fof(e9, plain, (sP1 <=> ((e1 = op(e3, e3)) & (e1 = op(e2, e2)) & (e1 = op(e1, e1)) & (e1 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f847, plain, (sP1 | sP0), inference(subsumption_resolution, [], [f846, f607])).
fof(f607, plain, ~ sP2, inference(subsumption_resolution, [], [f606, f349])).
fof(f349, plain, ~ (e2 = unit), inference(definition_unfolding, [], [f114, f135])).
fof(f114, plain, ~ (e0 = e2), inference(cnf_transformation, [], [f1])).
fof(f606, plain, ((e2 = unit) | ~ sP2), inference(forward_demodulation, [], [f335, f351])).
fof(f335, plain, ((e2 = op(e3, e3)) | ~ sP2), inference(cnf_transformation, [], [f110])).
fof(f110, plain, (((e2 = op(e3, e3)) & (e2 = op(e2, e2)) & (e2 = op(e1, e1)) & (e2 = op(e0, e0))) | ~ sP2), inference(nnf_transformation, [], [f10])).
fof(f10, plain, (((e2 = op(e3, e3)) & (e2 = op(e2, e2)) & (e2 = op(e1, e1)) & (e2 = op(e0, e0))) | ~ sP2), inference(usedef, [], [e10])).
fof(e10, plain, (sP2 <=> ((e2 = op(e3, e3)) & (e2 = op(e2, e2)) & (e2 = op(e1, e1)) & (e2 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f846, plain, (sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f845, f605])).
fof(f605, plain, ~ sP3, inference(subsumption_resolution, [], [f604, f348])).
fof(f604, plain, ((e3 = unit) | ~ sP3), inference(forward_demodulation, [], [f331, f351])).
fof(f331, plain, ((e3 = op(e3, e3)) | ~ sP3), inference(cnf_transformation, [], [f109])).
fof(f109, plain, (((e3 = op(e3, e3)) & (e3 = op(e2, e2)) & (e3 = op(e1, e1)) & (e3 = op(e0, e0))) | ~ sP3), inference(nnf_transformation, [], [f11])).
fof(f11, plain, (((e3 = op(e3, e3)) & (e3 = op(e2, e2)) & (e3 = op(e1, e1)) & (e3 = op(e0, e0))) | ~ sP3), inference(usedef, [], [e11])).
fof(e11, plain, (sP3 <=> ((e3 = op(e3, e3)) & (e3 = op(e2, e2)) & (e3 = op(e1, e1)) & (e3 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f845, plain, (sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f844, f603])).
fof(f603, plain, ~ sP4, inference(subsumption_resolution, [], [f471, f360])).
fof(f360, plain, (unit = op(unit, unit)), inference(definition_unfolding, [], [f119, f135, f135, f135])).
fof(f119, plain, (e0 = op(e0, e0)), inference(cnf_transformation, [], [f2])).
fof(f471, plain, (~ (unit = op(unit, unit)) | ~ sP4), inference(definition_unfolding, [], [f324, f135, f135, f135])).
fof(f324, plain, (~ (e0 = op(e0, e0)) | ~ sP4), inference(cnf_transformation, [], [f108])).
fof(f108, plain, ((~ (e3 = op(e0, e0)) & ~ (e2 = op(e0, e0)) & ~ (e1 = op(e0, e0)) & ~ (e0 = op(e0, e0))) | ~ sP4), inference(nnf_transformation, [], [f12])).
fof(f12, plain, ((~ (e3 = op(e0, e0)) & ~ (e2 = op(e0, e0)) & ~ (e1 = op(e0, e0)) & ~ (e0 = op(e0, e0))) | ~ sP4), inference(usedef, [], [e12])).
fof(e12, plain, (sP4 <=> (~ (e3 = op(e0, e0)) & ~ (e2 = op(e0, e0)) & ~ (e1 = op(e0, e0)) & ~ (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f844, plain, (sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f843, f599])).
fof(f599, plain, ~ sP5, inference(subsumption_resolution, [], [f466, f359])).
fof(f359, plain, (e1 = op(unit, e1)), inference(definition_unfolding, [], [f120, f135])).
fof(f120, plain, (e1 = op(e0, e1)), inference(cnf_transformation, [], [f2])).
fof(f466, plain, (~ (e1 = op(unit, e1)) | ~ sP5), inference(definition_unfolding, [], [f321, f135])).
fof(f321, plain, (~ (e1 = op(e0, e1)) | ~ sP5), inference(cnf_transformation, [], [f107])).
fof(f107, plain, ((~ (e3 = op(e0, e1)) & ~ (e2 = op(e0, e1)) & ~ (e1 = op(e0, e1)) & ~ (e0 = op(e0, e1))) | ~ sP5), inference(nnf_transformation, [], [f13])).
fof(f13, plain, ((~ (e3 = op(e0, e1)) & ~ (e2 = op(e0, e1)) & ~ (e1 = op(e0, e1)) & ~ (e0 = op(e0, e1))) | ~ sP5), inference(usedef, [], [e13])).
fof(e13, plain, (sP5 <=> (~ (e3 = op(e0, e1)) & ~ (e2 = op(e0, e1)) & ~ (e1 = op(e0, e1)) & ~ (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f843, plain, (sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f842, f596])).
fof(f596, plain, ~ sP6, inference(subsumption_resolution, [], [f461, f358])).
fof(f358, plain, (e2 = op(unit, e2)), inference(definition_unfolding, [], [f121, f135])).
fof(f121, plain, (e2 = op(e0, e2)), inference(cnf_transformation, [], [f2])).
fof(f461, plain, (~ (e2 = op(unit, e2)) | ~ sP6), inference(definition_unfolding, [], [f318, f135])).
fof(f318, plain, (~ (e2 = op(e0, e2)) | ~ sP6), inference(cnf_transformation, [], [f106])).
fof(f106, plain, ((~ (e3 = op(e0, e2)) & ~ (e2 = op(e0, e2)) & ~ (e1 = op(e0, e2)) & ~ (e0 = op(e0, e2))) | ~ sP6), inference(nnf_transformation, [], [f14])).
fof(f14, plain, ((~ (e3 = op(e0, e2)) & ~ (e2 = op(e0, e2)) & ~ (e1 = op(e0, e2)) & ~ (e0 = op(e0, e2))) | ~ sP6), inference(usedef, [], [e14])).
fof(e14, plain, (sP6 <=> (~ (e3 = op(e0, e2)) & ~ (e2 = op(e0, e2)) & ~ (e1 = op(e0, e2)) & ~ (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f842, plain, (sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f841, f594])).
fof(f594, plain, ~ sP7, inference(subsumption_resolution, [], [f456, f357])).
fof(f357, plain, (e3 = op(unit, e3)), inference(definition_unfolding, [], [f122, f135])).
fof(f122, plain, (e3 = op(e0, e3)), inference(cnf_transformation, [], [f2])).
fof(f456, plain, (~ (e3 = op(unit, e3)) | ~ sP7), inference(definition_unfolding, [], [f315, f135])).
fof(f315, plain, (~ (e3 = op(e0, e3)) | ~ sP7), inference(cnf_transformation, [], [f105])).
fof(f105, plain, ((~ (e3 = op(e0, e3)) & ~ (e2 = op(e0, e3)) & ~ (e1 = op(e0, e3)) & ~ (e0 = op(e0, e3))) | ~ sP7), inference(nnf_transformation, [], [f15])).
fof(f15, plain, ((~ (e3 = op(e0, e3)) & ~ (e2 = op(e0, e3)) & ~ (e1 = op(e0, e3)) & ~ (e0 = op(e0, e3))) | ~ sP7), inference(usedef, [], [e15])).
fof(e15, plain, (sP7 <=> (~ (e3 = op(e0, e3)) & ~ (e2 = op(e0, e3)) & ~ (e1 = op(e0, e3)) & ~ (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f841, plain, (sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f840, f593])).
fof(f593, plain, ~ sP8, inference(subsumption_resolution, [], [f454, f356])).
fof(f356, plain, (e1 = op(e1, unit)), inference(definition_unfolding, [], [f123, f135])).
fof(f123, plain, (e1 = op(e1, e0)), inference(cnf_transformation, [], [f2])).
fof(f454, plain, (~ (e1 = op(e1, unit)) | ~ sP8), inference(definition_unfolding, [], [f309, f135])).
fof(f309, plain, (~ (e1 = op(e1, e0)) | ~ sP8), inference(cnf_transformation, [], [f104])).
fof(f104, plain, ((~ (e3 = op(e1, e0)) & ~ (e2 = op(e1, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e1, e0))) | ~ sP8), inference(nnf_transformation, [], [f16])).
fof(f16, plain, ((~ (e3 = op(e1, e0)) & ~ (e2 = op(e1, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e1, e0))) | ~ sP8), inference(usedef, [], [e16])).
fof(e16, plain, (sP8 <=> (~ (e3 = op(e1, e0)) & ~ (e2 = op(e1, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f840, plain, (sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f839, f590])).
fof(f590, plain, ~ sP9, inference(subsumption_resolution, [], [f307, f124])).
fof(f124, plain, (e3 = op(e1, e1)), inference(cnf_transformation, [], [f2])).
fof(f307, plain, (~ (e3 = op(e1, e1)) | ~ sP9), inference(cnf_transformation, [], [f103])).
fof(f103, plain, ((~ (e3 = op(e1, e1)) & ~ (e2 = op(e1, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e1, e1))) | ~ sP9), inference(nnf_transformation, [], [f17])).
fof(f17, plain, ((~ (e3 = op(e1, e1)) & ~ (e2 = op(e1, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e1, e1))) | ~ sP9), inference(usedef, [], [e17])).
fof(e17, plain, (sP9 <=> (~ (e3 = op(e1, e1)) & ~ (e2 = op(e1, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f839, plain, (sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f838, f589])).
fof(f589, plain, ~ sP10, inference(subsumption_resolution, [], [f450, f355])).
fof(f355, plain, (op(e1, e2) = unit), inference(definition_unfolding, [], [f125, f135])).
fof(f125, plain, (e0 = op(e1, e2)), inference(cnf_transformation, [], [f2])).
fof(f450, plain, (~ (op(e1, e2) = unit) | ~ sP10), inference(definition_unfolding, [], [f300, f135])).
fof(f300, plain, (~ (e0 = op(e1, e2)) | ~ sP10), inference(cnf_transformation, [], [f102])).
fof(f102, plain, ((~ (e3 = op(e1, e2)) & ~ (e2 = op(e1, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e1, e2))) | ~ sP10), inference(nnf_transformation, [], [f18])).
fof(f18, plain, ((~ (e3 = op(e1, e2)) & ~ (e2 = op(e1, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e1, e2))) | ~ sP10), inference(usedef, [], [e18])).
fof(e18, plain, (sP10 <=> (~ (e3 = op(e1, e2)) & ~ (e2 = op(e1, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f838, plain, (sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f837, f585])).
fof(f585, plain, ~ sP11, inference(subsumption_resolution, [], [f298, f126])).
fof(f126, plain, (e2 = op(e1, e3)), inference(cnf_transformation, [], [f2])).
fof(f298, plain, (~ (e2 = op(e1, e3)) | ~ sP11), inference(cnf_transformation, [], [f101])).
fof(f101, plain, ((~ (e3 = op(e1, e3)) & ~ (e2 = op(e1, e3)) & ~ (e1 = op(e1, e3)) & ~ (e0 = op(e1, e3))) | ~ sP11), inference(nnf_transformation, [], [f19])).
fof(f19, plain, ((~ (e3 = op(e1, e3)) & ~ (e2 = op(e1, e3)) & ~ (e1 = op(e1, e3)) & ~ (e0 = op(e1, e3))) | ~ sP11), inference(usedef, [], [e19])).
fof(e19, plain, (sP11 <=> (~ (e3 = op(e1, e3)) & ~ (e2 = op(e1, e3)) & ~ (e1 = op(e1, e3)) & ~ (e0 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f837, plain, (sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f836, f583])).
fof(f583, plain, ~ sP12, inference(subsumption_resolution, [], [f446, f354])).
fof(f354, plain, (e2 = op(e2, unit)), inference(definition_unfolding, [], [f127, f135])).
fof(f127, plain, (e2 = op(e2, e0)), inference(cnf_transformation, [], [f2])).
fof(f446, plain, (~ (e2 = op(e2, unit)) | ~ sP12), inference(definition_unfolding, [], [f294, f135])).
fof(f294, plain, (~ (e2 = op(e2, e0)) | ~ sP12), inference(cnf_transformation, [], [f100])).
fof(f100, plain, ((~ (e3 = op(e2, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e2, e0)) & ~ (e0 = op(e2, e0))) | ~ sP12), inference(nnf_transformation, [], [f20])).
fof(f20, plain, ((~ (e3 = op(e2, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e2, e0)) & ~ (e0 = op(e2, e0))) | ~ sP12), inference(usedef, [], [e20])).
fof(e20, plain, (sP12 <=> (~ (e3 = op(e2, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e2, e0)) & ~ (e0 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f836, plain, (sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f835, f581])).
fof(f581, plain, ~ sP13, inference(subsumption_resolution, [], [f444, f353])).
fof(f353, plain, (op(e2, e1) = unit), inference(definition_unfolding, [], [f128, f135])).
fof(f128, plain, (e0 = op(e2, e1)), inference(cnf_transformation, [], [f2])).
fof(f444, plain, (~ (op(e2, e1) = unit) | ~ sP13), inference(definition_unfolding, [], [f288, f135])).
fof(f288, plain, (~ (e0 = op(e2, e1)) | ~ sP13), inference(cnf_transformation, [], [f99])).
fof(f99, plain, ((~ (e3 = op(e2, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e2, e1)) & ~ (e0 = op(e2, e1))) | ~ sP13), inference(nnf_transformation, [], [f21])).
fof(f21, plain, ((~ (e3 = op(e2, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e2, e1)) & ~ (e0 = op(e2, e1))) | ~ sP13), inference(usedef, [], [e21])).
fof(e21, plain, (sP13 <=> (~ (e3 = op(e2, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e2, e1)) & ~ (e0 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f835, plain, (sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f834, f577])).
fof(f577, plain, ~ sP14, inference(subsumption_resolution, [], [f287, f129])).
fof(f287, plain, (~ (e3 = op(e2, e2)) | ~ sP14), inference(cnf_transformation, [], [f98])).
fof(f98, plain, ((~ (e3 = op(e2, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e2, e2)) & ~ (e0 = op(e2, e2))) | ~ sP14), inference(nnf_transformation, [], [f22])).
fof(f22, plain, ((~ (e3 = op(e2, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e2, e2)) & ~ (e0 = op(e2, e2))) | ~ sP14), inference(usedef, [], [e22])).
fof(e22, plain, (sP14 <=> (~ (e3 = op(e2, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e2, e2)) & ~ (e0 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f834, plain, (sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f833, f576])).
fof(f576, plain, ~ sP15, inference(subsumption_resolution, [], [f281, f130])).
fof(f130, plain, (e1 = op(e2, e3)), inference(cnf_transformation, [], [f2])).
fof(f281, plain, (~ (e1 = op(e2, e3)) | ~ sP15), inference(cnf_transformation, [], [f97])).
fof(f97, plain, ((~ (e3 = op(e2, e3)) & ~ (e2 = op(e2, e3)) & ~ (e1 = op(e2, e3)) & ~ (e0 = op(e2, e3))) | ~ sP15), inference(nnf_transformation, [], [f23])).
fof(f23, plain, ((~ (e3 = op(e2, e3)) & ~ (e2 = op(e2, e3)) & ~ (e1 = op(e2, e3)) & ~ (e0 = op(e2, e3))) | ~ sP15), inference(usedef, [], [e23])).
fof(e23, plain, (sP15 <=> (~ (e3 = op(e2, e3)) & ~ (e2 = op(e2, e3)) & ~ (e1 = op(e2, e3)) & ~ (e0 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f833, plain, (sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f832, f573])).
fof(f573, plain, ~ sP16, inference(subsumption_resolution, [], [f438, f352])).
fof(f352, plain, (e3 = op(e3, unit)), inference(definition_unfolding, [], [f131, f135])).
fof(f131, plain, (e3 = op(e3, e0)), inference(cnf_transformation, [], [f2])).
fof(f438, plain, (~ (e3 = op(e3, unit)) | ~ sP16), inference(definition_unfolding, [], [f279, f135])).
fof(f279, plain, (~ (e3 = op(e3, e0)) | ~ sP16), inference(cnf_transformation, [], [f96])).
fof(f96, plain, ((~ (e3 = op(e3, e0)) & ~ (e2 = op(e3, e0)) & ~ (e1 = op(e3, e0)) & ~ (e0 = op(e3, e0))) | ~ sP16), inference(nnf_transformation, [], [f24])).
fof(f24, plain, ((~ (e3 = op(e3, e0)) & ~ (e2 = op(e3, e0)) & ~ (e1 = op(e3, e0)) & ~ (e0 = op(e3, e0))) | ~ sP16), inference(usedef, [], [e24])).
fof(e24, plain, (sP16 <=> (~ (e3 = op(e3, e0)) & ~ (e2 = op(e3, e0)) & ~ (e1 = op(e3, e0)) & ~ (e0 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f832, plain, (sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f831, f572])).
fof(f572, plain, ~ sP17, inference(subsumption_resolution, [], [f274, f132])).
fof(f132, plain, (e2 = op(e3, e1)), inference(cnf_transformation, [], [f2])).
fof(f274, plain, (~ (e2 = op(e3, e1)) | ~ sP17), inference(cnf_transformation, [], [f95])).
fof(f95, plain, ((~ (e3 = op(e3, e1)) & ~ (e2 = op(e3, e1)) & ~ (e1 = op(e3, e1)) & ~ (e0 = op(e3, e1))) | ~ sP17), inference(nnf_transformation, [], [f25])).
fof(f25, plain, ((~ (e3 = op(e3, e1)) & ~ (e2 = op(e3, e1)) & ~ (e1 = op(e3, e1)) & ~ (e0 = op(e3, e1))) | ~ sP17), inference(usedef, [], [e25])).
fof(e25, plain, (sP17 <=> (~ (e3 = op(e3, e1)) & ~ (e2 = op(e3, e1)) & ~ (e1 = op(e3, e1)) & ~ (e0 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f831, plain, (sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f830, f570])).
fof(f570, plain, ~ sP18, inference(subsumption_resolution, [], [f269, f133])).
fof(f133, plain, (e1 = op(e3, e2)), inference(cnf_transformation, [], [f2])).
fof(f269, plain, (~ (e1 = op(e3, e2)) | ~ sP18), inference(cnf_transformation, [], [f94])).
fof(f94, plain, ((~ (e3 = op(e3, e2)) & ~ (e2 = op(e3, e2)) & ~ (e1 = op(e3, e2)) & ~ (e0 = op(e3, e2))) | ~ sP18), inference(nnf_transformation, [], [f26])).
fof(f26, plain, ((~ (e3 = op(e3, e2)) & ~ (e2 = op(e3, e2)) & ~ (e1 = op(e3, e2)) & ~ (e0 = op(e3, e2))) | ~ sP18), inference(usedef, [], [e26])).
fof(e26, plain, (sP18 <=> (~ (e3 = op(e3, e2)) & ~ (e2 = op(e3, e2)) & ~ (e1 = op(e3, e2)) & ~ (e0 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f830, plain, (sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f829, f567])).
fof(f567, plain, ~ sP19, inference(subsumption_resolution, [], [f435, f351])).
fof(f435, plain, (~ (op(e3, e3) = unit) | ~ sP19), inference(definition_unfolding, [], [f264, f135])).
fof(f264, plain, (~ (e0 = op(e3, e3)) | ~ sP19), inference(cnf_transformation, [], [f93])).
fof(f93, plain, ((~ (e3 = op(e3, e3)) & ~ (e2 = op(e3, e3)) & ~ (e1 = op(e3, e3)) & ~ (e0 = op(e3, e3))) | ~ sP19), inference(nnf_transformation, [], [f27])).
fof(f27, plain, ((~ (e3 = op(e3, e3)) & ~ (e2 = op(e3, e3)) & ~ (e1 = op(e3, e3)) & ~ (e0 = op(e3, e3))) | ~ sP19), inference(usedef, [], [e27])).
fof(e27, plain, (sP19 <=> (~ (e3 = op(e3, e3)) & ~ (e2 = op(e3, e3)) & ~ (e1 = op(e3, e3)) & ~ (e0 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP19])])).
fof(f829, plain, (sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f828, f360])).
fof(f828, plain, (~ (unit = op(unit, unit)) | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f827, f359])).
fof(f827, plain, (~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f826, f356])).
fof(f826, plain, (~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f825, f358])).
fof(f825, plain, (~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f824, f354])).
fof(f824, plain, (~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f823, f352])).
fof(f823, plain, (~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f822, f487])).
fof(f487, plain, ~ sP20, inference(trivial_inequality_removal, [], [f434])).
fof(f434, plain, (~ (unit = unit) | ~ sP20), inference(definition_unfolding, [], [f260, f135])).
fof(f260, plain, (~ (e0 = unit) | ~ sP20), inference(cnf_transformation, [], [f92])).
fof(f92, plain, ((~ (e3 = unit) & ~ (e2 = unit) & ~ (e1 = unit) & ~ (e0 = unit)) | ~ sP20), inference(nnf_transformation, [], [f28])).
fof(f28, plain, ((~ (e3 = unit) & ~ (e2 = unit) & ~ (e1 = unit) & ~ (e0 = unit)) | ~ sP20), inference(usedef, [], [e28])).
fof(e28, plain, (sP20 <=> (~ (e3 = unit) & ~ (e2 = unit) & ~ (e1 = unit) & ~ (e0 = unit))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP20])])).
fof(f822, plain, (sP20 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f821, f563])).
fof(f563, plain, ~ sP21, inference(subsumption_resolution, [], [f433, f360])).
fof(f433, plain, (~ (unit = op(unit, unit)) | ~ sP21), inference(definition_unfolding, [], [f256, f135, f135, f135])).
fof(f256, plain, (~ (e0 = op(e0, e0)) | ~ sP21), inference(cnf_transformation, [], [f91])).
fof(f91, plain, ((~ (e0 = op(e0, e3)) & ~ (e0 = op(e0, e2)) & ~ (e0 = op(e0, e1)) & ~ (e0 = op(e0, e0))) | ~ sP21), inference(nnf_transformation, [], [f29])).
fof(f29, plain, ((~ (e0 = op(e0, e3)) & ~ (e0 = op(e0, e2)) & ~ (e0 = op(e0, e1)) & ~ (e0 = op(e0, e0))) | ~ sP21), inference(usedef, [], [e29])).
fof(e29, plain, (sP21 <=> (~ (e0 = op(e0, e3)) & ~ (e0 = op(e0, e2)) & ~ (e0 = op(e0, e1)) & ~ (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP21])])).
fof(f821, plain, (sP21 | sP20 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f820, f559])).
fof(f559, plain, ~ sP22, inference(subsumption_resolution, [], [f429, f360])).
fof(f429, plain, (~ (unit = op(unit, unit)) | ~ sP22), inference(definition_unfolding, [], [f252, f135, f135, f135])).
fof(f252, plain, (~ (e0 = op(e0, e0)) | ~ sP22), inference(cnf_transformation, [], [f90])).
fof(f90, plain, ((~ (e0 = op(e3, e0)) & ~ (e0 = op(e2, e0)) & ~ (e0 = op(e1, e0)) & ~ (e0 = op(e0, e0))) | ~ sP22), inference(nnf_transformation, [], [f30])).
fof(f30, plain, ((~ (e0 = op(e3, e0)) & ~ (e0 = op(e2, e0)) & ~ (e0 = op(e1, e0)) & ~ (e0 = op(e0, e0))) | ~ sP22), inference(usedef, [], [e30])).
fof(e30, plain, (sP22 <=> (~ (e0 = op(e3, e0)) & ~ (e0 = op(e2, e0)) & ~ (e0 = op(e1, e0)) & ~ (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP22])])).
fof(f820, plain, (sP22 | sP21 | sP20 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f819, f555])).
fof(f555, plain, ~ sP23, inference(subsumption_resolution, [], [f424, f359])).
fof(f424, plain, (~ (e1 = op(unit, e1)) | ~ sP23), inference(definition_unfolding, [], [f249, f135])).
fof(f249, plain, (~ (e1 = op(e0, e1)) | ~ sP23), inference(cnf_transformation, [], [f89])).
fof(f89, plain, ((~ (e1 = op(e0, e3)) & ~ (e1 = op(e0, e2)) & ~ (e1 = op(e0, e1)) & ~ (e1 = op(e0, e0))) | ~ sP23), inference(nnf_transformation, [], [f31])).
fof(f31, plain, ((~ (e1 = op(e0, e3)) & ~ (e1 = op(e0, e2)) & ~ (e1 = op(e0, e1)) & ~ (e1 = op(e0, e0))) | ~ sP23), inference(usedef, [], [e31])).
fof(e31, plain, (sP23 <=> (~ (e1 = op(e0, e3)) & ~ (e1 = op(e0, e2)) & ~ (e1 = op(e0, e1)) & ~ (e1 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP23])])).
fof(f819, plain, (sP23 | sP22 | sP21 | sP20 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f818, f552])).
fof(f552, plain, ~ sP24, inference(subsumption_resolution, [], [f420, f356])).
fof(f420, plain, (~ (e1 = op(e1, unit)) | ~ sP24), inference(definition_unfolding, [], [f245, f135])).
fof(f245, plain, (~ (e1 = op(e1, e0)) | ~ sP24), inference(cnf_transformation, [], [f88])).
fof(f88, plain, ((~ (e1 = op(e3, e0)) & ~ (e1 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e1 = op(e0, e0))) | ~ sP24), inference(nnf_transformation, [], [f32])).
fof(f32, plain, ((~ (e1 = op(e3, e0)) & ~ (e1 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e1 = op(e0, e0))) | ~ sP24), inference(usedef, [], [e32])).
fof(e32, plain, (sP24 <=> (~ (e1 = op(e3, e0)) & ~ (e1 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e1 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP24])])).
fof(f818, plain, (sP24 | sP23 | sP22 | sP21 | sP20 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f817, f549])).
fof(f549, plain, ~ sP25, inference(subsumption_resolution, [], [f415, f358])).
fof(f415, plain, (~ (e2 = op(unit, e2)) | ~ sP25), inference(definition_unfolding, [], [f242, f135])).
fof(f242, plain, (~ (e2 = op(e0, e2)) | ~ sP25), inference(cnf_transformation, [], [f87])).
fof(f87, plain, ((~ (e2 = op(e0, e3)) & ~ (e2 = op(e0, e2)) & ~ (e2 = op(e0, e1)) & ~ (e2 = op(e0, e0))) | ~ sP25), inference(nnf_transformation, [], [f33])).
fof(f33, plain, ((~ (e2 = op(e0, e3)) & ~ (e2 = op(e0, e2)) & ~ (e2 = op(e0, e1)) & ~ (e2 = op(e0, e0))) | ~ sP25), inference(usedef, [], [e33])).
fof(e33, plain, (sP25 <=> (~ (e2 = op(e0, e3)) & ~ (e2 = op(e0, e2)) & ~ (e2 = op(e0, e1)) & ~ (e2 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP25])])).
fof(f817, plain, (sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f816, f547])).
fof(f547, plain, ~ sP26, inference(subsumption_resolution, [], [f411, f354])).
fof(f411, plain, (~ (e2 = op(e2, unit)) | ~ sP26), inference(definition_unfolding, [], [f238, f135])).
fof(f238, plain, (~ (e2 = op(e2, e0)) | ~ sP26), inference(cnf_transformation, [], [f86])).
fof(f86, plain, ((~ (e2 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e2 = op(e1, e0)) & ~ (e2 = op(e0, e0))) | ~ sP26), inference(nnf_transformation, [], [f34])).
fof(f34, plain, ((~ (e2 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e2 = op(e1, e0)) & ~ (e2 = op(e0, e0))) | ~ sP26), inference(usedef, [], [e34])).
fof(e34, plain, (sP26 <=> (~ (e2 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e2 = op(e1, e0)) & ~ (e2 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP26])])).
fof(f816, plain, (sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f815, f545])).
fof(f545, plain, ~ sP27, inference(subsumption_resolution, [], [f406, f357])).
fof(f406, plain, (~ (e3 = op(unit, e3)) | ~ sP27), inference(definition_unfolding, [], [f235, f135])).
fof(f235, plain, (~ (e3 = op(e0, e3)) | ~ sP27), inference(cnf_transformation, [], [f85])).
fof(f85, plain, ((~ (e3 = op(e0, e3)) & ~ (e3 = op(e0, e2)) & ~ (e3 = op(e0, e1)) & ~ (e3 = op(e0, e0))) | ~ sP27), inference(nnf_transformation, [], [f35])).
fof(f35, plain, ((~ (e3 = op(e0, e3)) & ~ (e3 = op(e0, e2)) & ~ (e3 = op(e0, e1)) & ~ (e3 = op(e0, e0))) | ~ sP27), inference(usedef, [], [e35])).
fof(e35, plain, (sP27 <=> (~ (e3 = op(e0, e3)) & ~ (e3 = op(e0, e2)) & ~ (e3 = op(e0, e1)) & ~ (e3 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP27])])).
fof(f815, plain, (sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f814, f544])).
fof(f544, plain, ~ sP28, inference(subsumption_resolution, [], [f402, f352])).
fof(f402, plain, (~ (e3 = op(e3, unit)) | ~ sP28), inference(definition_unfolding, [], [f231, f135])).
fof(f231, plain, (~ (e3 = op(e3, e0)) | ~ sP28), inference(cnf_transformation, [], [f84])).
fof(f84, plain, ((~ (e3 = op(e3, e0)) & ~ (e3 = op(e2, e0)) & ~ (e3 = op(e1, e0)) & ~ (e3 = op(e0, e0))) | ~ sP28), inference(nnf_transformation, [], [f36])).
fof(f36, plain, ((~ (e3 = op(e3, e0)) & ~ (e3 = op(e2, e0)) & ~ (e3 = op(e1, e0)) & ~ (e3 = op(e0, e0))) | ~ sP28), inference(usedef, [], [e36])).
fof(e36, plain, (sP28 <=> (~ (e3 = op(e3, e0)) & ~ (e3 = op(e2, e0)) & ~ (e3 = op(e1, e0)) & ~ (e3 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP28])])).
fof(f814, plain, (sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f813, f543])).
fof(f543, plain, ~ sP29, inference(subsumption_resolution, [], [f399, f355])).
fof(f399, plain, (~ (op(e1, e2) = unit) | ~ sP29), inference(definition_unfolding, [], [f226, f135])).
fof(f226, plain, (~ (e0 = op(e1, e2)) | ~ sP29), inference(cnf_transformation, [], [f83])).
fof(f83, plain, ((~ (e0 = op(e1, e3)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e1, e0))) | ~ sP29), inference(nnf_transformation, [], [f37])).
fof(f37, plain, ((~ (e0 = op(e1, e3)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e1, e0))) | ~ sP29), inference(usedef, [], [e37])).
fof(e37, plain, (sP29 <=> (~ (e0 = op(e1, e3)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP29])])).
fof(f813, plain, (sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f812, f541])).
fof(f541, plain, ~ sP30, inference(subsumption_resolution, [], [f395, f353])).
fof(f395, plain, (~ (op(e2, e1) = unit) | ~ sP30), inference(definition_unfolding, [], [f222, f135])).
fof(f222, plain, (~ (e0 = op(e2, e1)) | ~ sP30), inference(cnf_transformation, [], [f82])).
fof(f82, plain, ((~ (e0 = op(e3, e1)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e0, e1))) | ~ sP30), inference(nnf_transformation, [], [f38])).
fof(f38, plain, ((~ (e0 = op(e3, e1)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e0, e1))) | ~ sP30), inference(usedef, [], [e38])).
fof(e38, plain, (sP30 <=> (~ (e0 = op(e3, e1)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP30])])).
fof(f812, plain, (sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f811, f539])).
fof(f539, plain, ~ sP31, inference(subsumption_resolution, [], [f393, f356])).
fof(f393, plain, (~ (e1 = op(e1, unit)) | ~ sP31), inference(definition_unfolding, [], [f216, f135])).
fof(f216, plain, (~ (e1 = op(e1, e0)) | ~ sP31), inference(cnf_transformation, [], [f81])).
fof(f81, plain, ((~ (e1 = op(e1, e3)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e1, e0))) | ~ sP31), inference(nnf_transformation, [], [f39])).
fof(f39, plain, ((~ (e1 = op(e1, e3)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e1, e0))) | ~ sP31), inference(usedef, [], [e39])).
fof(e39, plain, (sP31 <=> (~ (e1 = op(e1, e3)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP31])])).
fof(f811, plain, (sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f810, f535])).
fof(f535, plain, ~ sP32, inference(subsumption_resolution, [], [f392, f359])).
fof(f392, plain, (~ (e1 = op(unit, e1)) | ~ sP32), inference(definition_unfolding, [], [f212, f135])).
fof(f212, plain, (~ (e1 = op(e0, e1)) | ~ sP32), inference(cnf_transformation, [], [f80])).
fof(f80, plain, ((~ (e1 = op(e3, e1)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e0, e1))) | ~ sP32), inference(nnf_transformation, [], [f40])).
fof(f40, plain, ((~ (e1 = op(e3, e1)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e0, e1))) | ~ sP32), inference(usedef, [], [e40])).
fof(e40, plain, (sP32 <=> (~ (e1 = op(e3, e1)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP32])])).
fof(f810, plain, (sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f809, f531])).
fof(f531, plain, ~ sP33, inference(subsumption_resolution, [], [f211, f126])).
fof(f211, plain, (~ (e2 = op(e1, e3)) | ~ sP33), inference(cnf_transformation, [], [f79])).
fof(f79, plain, ((~ (e2 = op(e1, e3)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e1, e0))) | ~ sP33), inference(nnf_transformation, [], [f41])).
fof(f41, plain, ((~ (e2 = op(e1, e3)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e1, e0))) | ~ sP33), inference(usedef, [], [e41])).
fof(e41, plain, (sP33 <=> (~ (e2 = op(e1, e3)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP33])])).
fof(f809, plain, (sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f808, f530])).
fof(f530, plain, ~ sP34, inference(subsumption_resolution, [], [f207, f132])).
fof(f207, plain, (~ (e2 = op(e3, e1)) | ~ sP34), inference(cnf_transformation, [], [f78])).
fof(f78, plain, ((~ (e2 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e0, e1))) | ~ sP34), inference(nnf_transformation, [], [f42])).
fof(f42, plain, ((~ (e2 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e0, e1))) | ~ sP34), inference(usedef, [], [e42])).
fof(e42, plain, (sP34 <=> (~ (e2 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP34])])).
fof(f808, plain, (sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f807, f529])).
fof(f529, plain, ~ sP35, inference(subsumption_resolution, [], [f201, f124])).
fof(f201, plain, (~ (e3 = op(e1, e1)) | ~ sP35), inference(cnf_transformation, [], [f77])).
fof(f77, plain, ((~ (e3 = op(e1, e3)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e1, e0))) | ~ sP35), inference(nnf_transformation, [], [f43])).
fof(f43, plain, ((~ (e3 = op(e1, e3)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e1, e0))) | ~ sP35), inference(usedef, [], [e43])).
fof(e43, plain, (sP35 <=> (~ (e3 = op(e1, e3)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP35])])).
fof(f807, plain, (sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f806, f526])).
fof(f526, plain, ~ sP36, inference(subsumption_resolution, [], [f197, f124])).
fof(f197, plain, (~ (e3 = op(e1, e1)) | ~ sP36), inference(cnf_transformation, [], [f76])).
fof(f76, plain, ((~ (e3 = op(e3, e1)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e0, e1))) | ~ sP36), inference(nnf_transformation, [], [f44])).
fof(f44, plain, ((~ (e3 = op(e3, e1)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e0, e1))) | ~ sP36), inference(usedef, [], [e44])).
fof(e44, plain, (sP36 <=> (~ (e3 = op(e3, e1)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP36])])).
fof(f806, plain, (sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f805, f523])).
fof(f523, plain, ~ sP37, inference(subsumption_resolution, [], [f386, f353])).
fof(f386, plain, (~ (op(e2, e1) = unit) | ~ sP37), inference(definition_unfolding, [], [f193, f135])).
fof(f193, plain, (~ (e0 = op(e2, e1)) | ~ sP37), inference(cnf_transformation, [], [f75])).
fof(f75, plain, ((~ (e0 = op(e2, e3)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e2, e0))) | ~ sP37), inference(nnf_transformation, [], [f45])).
fof(f45, plain, ((~ (e0 = op(e2, e3)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e2, e0))) | ~ sP37), inference(usedef, [], [e45])).
fof(e45, plain, (sP37 <=> (~ (e0 = op(e2, e3)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP37])])).
fof(f805, plain, (sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f804, f520])).
fof(f520, plain, ~ sP38, inference(subsumption_resolution, [], [f382, f355])).
fof(f382, plain, (~ (op(e1, e2) = unit) | ~ sP38), inference(definition_unfolding, [], [f189, f135])).
fof(f189, plain, (~ (e0 = op(e1, e2)) | ~ sP38), inference(cnf_transformation, [], [f74])).
fof(f74, plain, ((~ (e0 = op(e3, e2)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e0, e2))) | ~ sP38), inference(nnf_transformation, [], [f46])).
fof(f46, plain, ((~ (e0 = op(e3, e2)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e0, e2))) | ~ sP38), inference(usedef, [], [e46])).
fof(e46, plain, (sP38 <=> (~ (e0 = op(e3, e2)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP38])])).
fof(f804, plain, (sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f803, f517])).
fof(f517, plain, ~ sP39, inference(subsumption_resolution, [], [f187, f130])).
fof(f187, plain, (~ (e1 = op(e2, e3)) | ~ sP39), inference(cnf_transformation, [], [f73])).
fof(f73, plain, ((~ (e1 = op(e2, e3)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e2, e0))) | ~ sP39), inference(nnf_transformation, [], [f47])).
fof(f47, plain, ((~ (e1 = op(e2, e3)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e2, e0))) | ~ sP39), inference(usedef, [], [e47])).
fof(e47, plain, (sP39 <=> (~ (e1 = op(e2, e3)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP39])])).
fof(f803, plain, (sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f802, f516])).
fof(f516, plain, ~ sP40, inference(subsumption_resolution, [], [f183, f133])).
fof(f183, plain, (~ (e1 = op(e3, e2)) | ~ sP40), inference(cnf_transformation, [], [f72])).
fof(f72, plain, ((~ (e1 = op(e3, e2)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e0, e2))) | ~ sP40), inference(nnf_transformation, [], [f48])).
fof(f48, plain, ((~ (e1 = op(e3, e2)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e0, e2))) | ~ sP40), inference(usedef, [], [e48])).
fof(e48, plain, (sP40 <=> (~ (e1 = op(e3, e2)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP40])])).
fof(f802, plain, (sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f801, f515])).
fof(f515, plain, ~ sP41, inference(subsumption_resolution, [], [f377, f354])).
fof(f377, plain, (~ (e2 = op(e2, unit)) | ~ sP41), inference(definition_unfolding, [], [f176, f135])).
fof(f176, plain, (~ (e2 = op(e2, e0)) | ~ sP41), inference(cnf_transformation, [], [f71])).
fof(f71, plain, ((~ (e2 = op(e2, e3)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e2, e0))) | ~ sP41), inference(nnf_transformation, [], [f49])).
fof(f49, plain, ((~ (e2 = op(e2, e3)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e2, e0))) | ~ sP41), inference(usedef, [], [e49])).
fof(e49, plain, (sP41 <=> (~ (e2 = op(e2, e3)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP41])])).
fof(f801, plain, (sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f800, f511])).
fof(f511, plain, ~ sP42, inference(subsumption_resolution, [], [f376, f358])).
fof(f376, plain, (~ (e2 = op(unit, e2)) | ~ sP42), inference(definition_unfolding, [], [f172, f135])).
fof(f172, plain, (~ (e2 = op(e0, e2)) | ~ sP42), inference(cnf_transformation, [], [f70])).
fof(f70, plain, ((~ (e2 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e0, e2))) | ~ sP42), inference(nnf_transformation, [], [f50])).
fof(f50, plain, ((~ (e2 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e0, e2))) | ~ sP42), inference(usedef, [], [e50])).
fof(e50, plain, (sP42 <=> (~ (e2 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP42])])).
fof(f800, plain, (sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f799, f507])).
fof(f507, plain, ~ sP43, inference(subsumption_resolution, [], [f170, f129])).
fof(f170, plain, (~ (e3 = op(e2, e2)) | ~ sP43), inference(cnf_transformation, [], [f69])).
fof(f69, plain, ((~ (e3 = op(e2, e3)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e2, e0))) | ~ sP43), inference(nnf_transformation, [], [f51])).
fof(f51, plain, ((~ (e3 = op(e2, e3)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e2, e0))) | ~ sP43), inference(usedef, [], [e51])).
fof(e51, plain, (sP43 <=> (~ (e3 = op(e2, e3)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP43])])).
fof(f799, plain, (sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f798, f505])).
fof(f505, plain, ~ sP44, inference(subsumption_resolution, [], [f166, f129])).
fof(f166, plain, (~ (e3 = op(e2, e2)) | ~ sP44), inference(cnf_transformation, [], [f68])).
fof(f68, plain, ((~ (e3 = op(e3, e2)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e0, e2))) | ~ sP44), inference(nnf_transformation, [], [f52])).
fof(f52, plain, ((~ (e3 = op(e3, e2)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e0, e2))) | ~ sP44), inference(usedef, [], [e52])).
fof(e52, plain, (sP44 <=> (~ (e3 = op(e3, e2)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP44])])).
fof(f798, plain, (sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f797, f503])).
fof(f503, plain, ~ sP45, inference(subsumption_resolution, [], [f370, f351])).
fof(f370, plain, (~ (op(e3, e3) = unit) | ~ sP45), inference(definition_unfolding, [], [f163, f135])).
fof(f163, plain, (~ (e0 = op(e3, e3)) | ~ sP45), inference(cnf_transformation, [], [f67])).
fof(f67, plain, ((~ (e0 = op(e3, e3)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e3, e0))) | ~ sP45), inference(nnf_transformation, [], [f53])).
fof(f53, plain, ((~ (e0 = op(e3, e3)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e3, e0))) | ~ sP45), inference(usedef, [], [e53])).
fof(e53, plain, (sP45 <=> (~ (e0 = op(e3, e3)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP45])])).
fof(f797, plain, (sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f796, f502])).
fof(f502, plain, ~ sP46, inference(subsumption_resolution, [], [f366, f351])).
fof(f366, plain, (~ (op(e3, e3) = unit) | ~ sP46), inference(definition_unfolding, [], [f159, f135])).
fof(f159, plain, (~ (e0 = op(e3, e3)) | ~ sP46), inference(cnf_transformation, [], [f66])).
fof(f66, plain, ((~ (e0 = op(e3, e3)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e0, e3))) | ~ sP46), inference(nnf_transformation, [], [f54])).
fof(f54, plain, ((~ (e0 = op(e3, e3)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e0, e3))) | ~ sP46), inference(usedef, [], [e54])).
fof(e54, plain, (sP46 <=> (~ (e0 = op(e3, e3)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP46])])).
fof(f796, plain, (sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f795, f501])).
fof(f501, plain, ~ sP47, inference(subsumption_resolution, [], [f154, f133])).
fof(f154, plain, (~ (e1 = op(e3, e2)) | ~ sP47), inference(cnf_transformation, [], [f65])).
fof(f65, plain, ((~ (e1 = op(e3, e3)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e3, e0))) | ~ sP47), inference(nnf_transformation, [], [f55])).
fof(f55, plain, ((~ (e1 = op(e3, e3)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e3, e0))) | ~ sP47), inference(usedef, [], [e55])).
fof(e55, plain, (sP47 <=> (~ (e1 = op(e3, e3)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP47])])).
fof(f795, plain, (sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f794, f499])).
fof(f499, plain, ~ sP48, inference(subsumption_resolution, [], [f150, f130])).
fof(f150, plain, (~ (e1 = op(e2, e3)) | ~ sP48), inference(cnf_transformation, [], [f64])).
fof(f64, plain, ((~ (e1 = op(e3, e3)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e0, e3))) | ~ sP48), inference(nnf_transformation, [], [f56])).
fof(f56, plain, ((~ (e1 = op(e3, e3)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e0, e3))) | ~ sP48), inference(usedef, [], [e56])).
fof(e56, plain, (sP48 <=> (~ (e1 = op(e3, e3)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP48])])).
fof(f794, plain, (sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f793, f497])).
fof(f497, plain, ~ sP49, inference(subsumption_resolution, [], [f145, f132])).
fof(f145, plain, (~ (e2 = op(e3, e1)) | ~ sP49), inference(cnf_transformation, [], [f63])).
fof(f63, plain, ((~ (e2 = op(e3, e3)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e3, e0))) | ~ sP49), inference(nnf_transformation, [], [f57])).
fof(f57, plain, ((~ (e2 = op(e3, e3)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e3, e0))) | ~ sP49), inference(usedef, [], [e57])).
fof(e57, plain, (sP49 <=> (~ (e2 = op(e3, e3)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP49])])).
fof(f793, plain, (sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f792, f494])).
fof(f494, plain, ~ sP50, inference(subsumption_resolution, [], [f141, f126])).
fof(f141, plain, (~ (e2 = op(e1, e3)) | ~ sP50), inference(cnf_transformation, [], [f62])).
fof(f62, plain, ((~ (e2 = op(e3, e3)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e0, e3))) | ~ sP50), inference(nnf_transformation, [], [f58])).
fof(f58, plain, ((~ (e2 = op(e3, e3)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e0, e3))) | ~ sP50), inference(usedef, [], [e58])).
fof(e58, plain, (sP50 <=> (~ (e2 = op(e3, e3)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP50])])).
fof(f792, plain, (sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f791, f491])).
fof(f491, plain, ~ sP51, inference(subsumption_resolution, [], [f361, f352])).
fof(f361, plain, (~ (e3 = op(e3, unit)) | ~ sP51), inference(definition_unfolding, [], [f136, f135])).
fof(f136, plain, (~ (e3 = op(e3, e0)) | ~ sP51), inference(cnf_transformation, [], [f61])).
fof(f61, plain, ((~ (e3 = op(e3, e3)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e3, e0))) | ~ sP51), inference(nnf_transformation, [], [f59])).
fof(f59, plain, ((~ (e3 = op(e3, e3)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e3, e0))) | ~ sP51), inference(usedef, [], [e59])).
fof(e59, plain, (sP51 <=> (~ (e3 = op(e3, e3)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP51])])).
fof(f791, plain, (sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f483, f357])).
fof(f483, plain, (~ (e3 = op(unit, e3)) | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(duplicate_literal_removal, [], [f482])).
fof(f482, plain, (~ (e3 = op(unit, e3)) | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | ~ (unit = op(unit, unit)) | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(definition_unfolding, [], [f344, f135, f135, f135, f135, f135])).
fof(f344, plain, (~ (e3 = op(e0, e3)) | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (e0 = op(e0, unit)) | ~ (e0 = op(unit, e0)) | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f60])).
fof(f60, plain, ((~ (e3 = op(e3, e3)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e0, e3))) | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (e0 = op(e0, unit)) | ~ (e0 = op(unit, e0)) | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(definition_folding, [], [f7, e59, e58, e57, e56, e55, e54, e53, e52, e51, e50, e49, e48, e47, e46, e45, e44, e43, e42, e41, e40, e39, e38, e37, e36, e35, e34, e33, e32, e31, e30, e29, e28, e27, e26, e25, e24, e23, e22, e21, e20, e19, e18, e17, e16, e15, e14, e13, e12, e11, e10, e9, e8])).
fof(f7, plain, ((~ (e3 = op(e3, e3)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e0, e3))) | (~ (e3 = op(e3, e3)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e3, e0))) | (~ (e2 = op(e3, e3)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e0, e3))) | (~ (e2 = op(e3, e3)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e3, e0))) | (~ (e1 = op(e3, e3)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e0, e3))) | (~ (e1 = op(e3, e3)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e3, e0))) | (~ (e0 = op(e3, e3)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e0, e3))) | (~ (e0 = op(e3, e3)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e3, e0))) | (~ (e3 = op(e3, e2)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e0, e2))) | (~ (e3 = op(e2, e3)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e2, e0))) | (~ (e2 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e0, e2))) | (~ (e2 = op(e2, e3)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e2, e0))) | (~ (e1 = op(e3, e2)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e0, e2))) | (~ (e1 = op(e2, e3)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e2, e0))) | (~ (e0 = op(e3, e2)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e0, e2))) | (~ (e0 = op(e2, e3)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e2, e0))) | (~ (e3 = op(e3, e1)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e0, e1))) | (~ (e3 = op(e1, e3)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e1, e0))) | (~ (e2 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e0, e1))) | (~ (e2 = op(e1, e3)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e1, e0))) | (~ (e1 = op(e3, e1)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e0, e1))) | (~ (e1 = op(e1, e3)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e1, e0))) | (~ (e0 = op(e3, e1)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e0, e1))) | (~ (e0 = op(e1, e3)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e1, e0))) | (~ (e3 = op(e3, e0)) & ~ (e3 = op(e2, e0)) & ~ (e3 = op(e1, e0)) & ~ (e3 = op(e0, e0))) | (~ (e3 = op(e0, e3)) & ~ (e3 = op(e0, e2)) & ~ (e3 = op(e0, e1)) & ~ (e3 = op(e0, e0))) | (~ (e2 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e2 = op(e1, e0)) & ~ (e2 = op(e0, e0))) | (~ (e2 = op(e0, e3)) & ~ (e2 = op(e0, e2)) & ~ (e2 = op(e0, e1)) & ~ (e2 = op(e0, e0))) | (~ (e1 = op(e3, e0)) & ~ (e1 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e1 = op(e0, e0))) | (~ (e1 = op(e0, e3)) & ~ (e1 = op(e0, e2)) & ~ (e1 = op(e0, e1)) & ~ (e1 = op(e0, e0))) | (~ (e0 = op(e3, e0)) & ~ (e0 = op(e2, e0)) & ~ (e0 = op(e1, e0)) & ~ (e0 = op(e0, e0))) | (~ (e0 = op(e0, e3)) & ~ (e0 = op(e0, e2)) & ~ (e0 = op(e0, e1)) & ~ (e0 = op(e0, e0))) | (~ (e3 = unit) & ~ (e2 = unit) & ~ (e1 = unit) & ~ (e0 = unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (e0 = op(e0, unit)) | ~ (e0 = op(unit, e0)) | (~ (e3 = op(e3, e3)) & ~ (e2 = op(e3, e3)) & ~ (e1 = op(e3, e3)) & ~ (e0 = op(e3, e3))) | (~ (e3 = op(e3, e2)) & ~ (e2 = op(e3, e2)) & ~ (e1 = op(e3, e2)) & ~ (e0 = op(e3, e2))) | (~ (e3 = op(e3, e1)) & ~ (e2 = op(e3, e1)) & ~ (e1 = op(e3, e1)) & ~ (e0 = op(e3, e1))) | (~ (e3 = op(e3, e0)) & ~ (e2 = op(e3, e0)) & ~ (e1 = op(e3, e0)) & ~ (e0 = op(e3, e0))) | (~ (e3 = op(e2, e3)) & ~ (e2 = op(e2, e3)) & ~ (e1 = op(e2, e3)) & ~ (e0 = op(e2, e3))) | (~ (e3 = op(e2, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e2, e2)) & ~ (e0 = op(e2, e2))) | (~ (e3 = op(e2, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e2, e1)) & ~ (e0 = op(e2, e1))) | (~ (e3 = op(e2, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e2, e0)) & ~ (e0 = op(e2, e0))) | (~ (e3 = op(e1, e3)) & ~ (e2 = op(e1, e3)) & ~ (e1 = op(e1, e3)) & ~ (e0 = op(e1, e3))) | (~ (e3 = op(e1, e2)) & ~ (e2 = op(e1, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e1, e2))) | (~ (e3 = op(e1, e1)) & ~ (e2 = op(e1, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e1, e1))) | (~ (e3 = op(e1, e0)) & ~ (e2 = op(e1, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e1, e0))) | (~ (e3 = op(e0, e3)) & ~ (e2 = op(e0, e3)) & ~ (e1 = op(e0, e3)) & ~ (e0 = op(e0, e3))) | (~ (e3 = op(e0, e2)) & ~ (e2 = op(e0, e2)) & ~ (e1 = op(e0, e2)) & ~ (e0 = op(e0, e2))) | (~ (e3 = op(e0, e1)) & ~ (e2 = op(e0, e1)) & ~ (e1 = op(e0, e1)) & ~ (e0 = op(e0, e1))) | (~ (e3 = op(e0, e0)) & ~ (e2 = op(e0, e0)) & ~ (e1 = op(e0, e0)) & ~ (e0 = op(e0, e0))) | ((e3 = op(e3, e3)) & (e3 = op(e2, e2)) & (e3 = op(e1, e1)) & (e3 = op(e0, e0))) | ((e2 = op(e3, e3)) & (e2 = op(e2, e2)) & (e2 = op(e1, e1)) & (e2 = op(e0, e0))) | ((e1 = op(e3, e3)) & (e1 = op(e2, e2)) & (e1 = op(e1, e1)) & (e1 = op(e0, e0))) | ((e0 = op(e3, e3)) & (e0 = op(e2, e2)) & (e0 = op(e1, e1)) & (e0 = op(e0, e0)))), inference(flattening, [], [f6])).
fof(f6, plain, ((~ (e3 = op(e3, e3)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e0, e3))) | (~ (e3 = op(e3, e3)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e3, e0))) | (~ (e2 = op(e3, e3)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e0, e3))) | (~ (e2 = op(e3, e3)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e3, e0))) | (~ (e1 = op(e3, e3)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e0, e3))) | (~ (e1 = op(e3, e3)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e3, e0))) | (~ (e0 = op(e3, e3)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e0, e3))) | (~ (e0 = op(e3, e3)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e3, e0))) | (~ (e3 = op(e3, e2)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e0, e2))) | (~ (e3 = op(e2, e3)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e2, e0))) | (~ (e2 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e0, e2))) | (~ (e2 = op(e2, e3)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e2, e0))) | (~ (e1 = op(e3, e2)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e0, e2))) | (~ (e1 = op(e2, e3)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e2, e0))) | (~ (e0 = op(e3, e2)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e0, e2))) | (~ (e0 = op(e2, e3)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e2, e0))) | (~ (e3 = op(e3, e1)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e0, e1))) | (~ (e3 = op(e1, e3)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e1, e0))) | (~ (e2 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e0, e1))) | (~ (e2 = op(e1, e3)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e1, e0))) | (~ (e1 = op(e3, e1)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e0, e1))) | (~ (e1 = op(e1, e3)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e1, e0))) | (~ (e0 = op(e3, e1)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e0, e1))) | (~ (e0 = op(e1, e3)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e1, e0))) | (~ (e3 = op(e3, e0)) & ~ (e3 = op(e2, e0)) & ~ (e3 = op(e1, e0)) & ~ (e3 = op(e0, e0))) | (~ (e3 = op(e0, e3)) & ~ (e3 = op(e0, e2)) & ~ (e3 = op(e0, e1)) & ~ (e3 = op(e0, e0))) | (~ (e2 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e2 = op(e1, e0)) & ~ (e2 = op(e0, e0))) | (~ (e2 = op(e0, e3)) & ~ (e2 = op(e0, e2)) & ~ (e2 = op(e0, e1)) & ~ (e2 = op(e0, e0))) | (~ (e1 = op(e3, e0)) & ~ (e1 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e1 = op(e0, e0))) | (~ (e1 = op(e0, e3)) & ~ (e1 = op(e0, e2)) & ~ (e1 = op(e0, e1)) & ~ (e1 = op(e0, e0))) | (~ (e0 = op(e3, e0)) & ~ (e0 = op(e2, e0)) & ~ (e0 = op(e1, e0)) & ~ (e0 = op(e0, e0))) | (~ (e0 = op(e0, e3)) & ~ (e0 = op(e0, e2)) & ~ (e0 = op(e0, e1)) & ~ (e0 = op(e0, e0))) | (~ (e3 = unit) & ~ (e2 = unit) & ~ (e1 = unit) & ~ (e0 = unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (e0 = op(e0, unit)) | ~ (e0 = op(unit, e0)) | (~ (e3 = op(e3, e3)) & ~ (e2 = op(e3, e3)) & ~ (e1 = op(e3, e3)) & ~ (e0 = op(e3, e3))) | (~ (e3 = op(e3, e2)) & ~ (e2 = op(e3, e2)) & ~ (e1 = op(e3, e2)) & ~ (e0 = op(e3, e2))) | (~ (e3 = op(e3, e1)) & ~ (e2 = op(e3, e1)) & ~ (e1 = op(e3, e1)) & ~ (e0 = op(e3, e1))) | (~ (e3 = op(e3, e0)) & ~ (e2 = op(e3, e0)) & ~ (e1 = op(e3, e0)) & ~ (e0 = op(e3, e0))) | (~ (e3 = op(e2, e3)) & ~ (e2 = op(e2, e3)) & ~ (e1 = op(e2, e3)) & ~ (e0 = op(e2, e3))) | (~ (e3 = op(e2, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e2, e2)) & ~ (e0 = op(e2, e2))) | (~ (e3 = op(e2, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e2, e1)) & ~ (e0 = op(e2, e1))) | (~ (e3 = op(e2, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e2, e0)) & ~ (e0 = op(e2, e0))) | (~ (e3 = op(e1, e3)) & ~ (e2 = op(e1, e3)) & ~ (e1 = op(e1, e3)) & ~ (e0 = op(e1, e3))) | (~ (e3 = op(e1, e2)) & ~ (e2 = op(e1, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e1, e2))) | (~ (e3 = op(e1, e1)) & ~ (e2 = op(e1, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e1, e1))) | (~ (e3 = op(e1, e0)) & ~ (e2 = op(e1, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e1, e0))) | (~ (e3 = op(e0, e3)) & ~ (e2 = op(e0, e3)) & ~ (e1 = op(e0, e3)) & ~ (e0 = op(e0, e3))) | (~ (e3 = op(e0, e2)) & ~ (e2 = op(e0, e2)) & ~ (e1 = op(e0, e2)) & ~ (e0 = op(e0, e2))) | (~ (e3 = op(e0, e1)) & ~ (e2 = op(e0, e1)) & ~ (e1 = op(e0, e1)) & ~ (e0 = op(e0, e1))) | (~ (e3 = op(e0, e0)) & ~ (e2 = op(e0, e0)) & ~ (e1 = op(e0, e0)) & ~ (e0 = op(e0, e0))) | (((e3 = op(e3, e3)) & (e3 = op(e2, e2)) & (e3 = op(e1, e1)) & (e3 = op(e0, e0))) | ((e2 = op(e3, e3)) & (e2 = op(e2, e2)) & (e2 = op(e1, e1)) & (e2 = op(e0, e0))) | ((e1 = op(e3, e3)) & (e1 = op(e2, e2)) & (e1 = op(e1, e1)) & (e1 = op(e0, e0))) | ((e0 = op(e3, e3)) & (e0 = op(e2, e2)) & (e0 = op(e1, e1)) & (e0 = op(e0, e0))))), inference(ennf_transformation, [], [f5])).
fof(f5, plain, ~ (((e3 = op(e3, e3)) | (e3 = op(e2, e3)) | (e3 = op(e1, e3)) | (e3 = op(e0, e3))) & ((e3 = op(e3, e3)) | (e3 = op(e3, e2)) | (e3 = op(e3, e1)) | (e3 = op(e3, e0))) & ((e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(e0, e3))) & ((e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, e0))) & ((e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(e0, e3))) & ((e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, e0))) & ((e0 = op(e3, e3)) | (e0 = op(e2, e3)) | (e0 = op(e1, e3)) | (e0 = op(e0, e3))) & ((e0 = op(e3, e3)) | (e0 = op(e3, e2)) | (e0 = op(e3, e1)) | (e0 = op(e3, e0))) & ((e3 = op(e3, e2)) | (e3 = op(e2, e2)) | (e3 = op(e1, e2)) | (e3 = op(e0, e2))) & ((e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, e0))) & ((e2 = op(e3, e2)) | (e2 = op(e2, e2)) | (e2 = op(e1, e2)) | (e2 = op(e0, e2))) & ((e2 = op(e2, e3)) | (e2 = op(e2, e2)) | (e2 = op(e2, e1)) | (e2 = op(e2, e0))) & ((e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(e0, e2))) & ((e1 = op(e2, e3)) | (e1 = op(e2, e2)) | (e1 = op(e2, e1)) | (e1 = op(e2, e0))) & ((e0 = op(e3, e2)) | (e0 = op(e2, e2)) | (e0 = op(e1, e2)) | (e0 = op(e0, e2))) & ((e0 = op(e2, e3)) | (e0 = op(e2, e2)) | (e0 = op(e2, e1)) | (e0 = op(e2, e0))) & ((e3 = op(e3, e1)) | (e3 = op(e2, e1)) | (e3 = op(e1, e1)) | (e3 = op(e0, e1))) & ((e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, e0))) & ((e2 = op(e3, e1)) | (e2 = op(e2, e1)) | (e2 = op(e1, e1)) | (e2 = op(e0, e1))) & ((e2 = op(e1, e3)) | (e2 = op(e1, e2)) | (e2 = op(e1, e1)) | (e2 = op(e1, e0))) & ((e1 = op(e3, e1)) | (e1 = op(e2, e1)) | (e1 = op(e1, e1)) | (e1 = op(e0, e1))) & ((e1 = op(e1, e3)) | (e1 = op(e1, e2)) | (e1 = op(e1, e1)) | (e1 = op(e1, e0))) & ((e0 = op(e3, e1)) | (e0 = op(e2, e1)) | (e0 = op(e1, e1)) | (e0 = op(e0, e1))) & ((e0 = op(e1, e3)) | (e0 = op(e1, e2)) | (e0 = op(e1, e1)) | (e0 = op(e1, e0))) & ((e3 = op(e3, e0)) | (e3 = op(e2, e0)) | (e3 = op(e1, e0)) | (e3 = op(e0, e0))) & ((e3 = op(e0, e3)) | (e3 = op(e0, e2)) | (e3 = op(e0, e1)) | (e3 = op(e0, e0))) & ((e2 = op(e3, e0)) | (e2 = op(e2, e0)) | (e2 = op(e1, e0)) | (e2 = op(e0, e0))) & ((e2 = op(e0, e3)) | (e2 = op(e0, e2)) | (e2 = op(e0, e1)) | (e2 = op(e0, e0))) & ((e1 = op(e3, e0)) | (e1 = op(e2, e0)) | (e1 = op(e1, e0)) | (e1 = op(e0, e0))) & ((e1 = op(e0, e3)) | (e1 = op(e0, e2)) | (e1 = op(e0, e1)) | (e1 = op(e0, e0))) & ((e0 = op(e3, e0)) | (e0 = op(e2, e0)) | (e0 = op(e1, e0)) | (e0 = op(e0, e0))) & ((e0 = op(e0, e3)) | (e0 = op(e0, e2)) | (e0 = op(e0, e1)) | (e0 = op(e0, e0))) & ((e3 = unit) | (e2 = unit) | (e1 = unit) | (e0 = unit)) & (e3 = op(e3, unit)) & (e3 = op(unit, e3)) & (e2 = op(e2, unit)) & (e2 = op(unit, e2)) & (e1 = op(e1, unit)) & (e1 = op(unit, e1)) & (e0 = op(e0, unit)) & (e0 = op(unit, e0)) & ((e3 = op(e3, e3)) | (e2 = op(e3, e3)) | (e1 = op(e3, e3)) | (e0 = op(e3, e3))) & ((e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (e0 = op(e3, e2))) & ((e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (e0 = op(e3, e1))) & ((e3 = op(e3, e0)) | (e2 = op(e3, e0)) | (e1 = op(e3, e0)) | (e0 = op(e3, e0))) & ((e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))) & ((e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (e0 = op(e2, e2))) & ((e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))) & ((e3 = op(e2, e0)) | (e2 = op(e2, e0)) | (e1 = op(e2, e0)) | (e0 = op(e2, e0))) & ((e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (e0 = op(e1, e3))) & ((e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (e0 = op(e1, e2))) & ((e3 = op(e1, e1)) | (e2 = op(e1, e1)) | (e1 = op(e1, e1)) | (e0 = op(e1, e1))) & ((e3 = op(e1, e0)) | (e2 = op(e1, e0)) | (e1 = op(e1, e0)) | (e0 = op(e1, e0))) & ((e3 = op(e0, e3)) | (e2 = op(e0, e3)) | (e1 = op(e0, e3)) | (e0 = op(e0, e3))) & ((e3 = op(e0, e2)) | (e2 = op(e0, e2)) | (e1 = op(e0, e2)) | (e0 = op(e0, e2))) & ((e3 = op(e0, e1)) | (e2 = op(e0, e1)) | (e1 = op(e0, e1)) | (e0 = op(e0, e1))) & ((e3 = op(e0, e0)) | (e2 = op(e0, e0)) | (e1 = op(e0, e0)) | (e0 = op(e0, e0))) & ~ (((e3 = op(e3, e3)) & (e3 = op(e2, e2)) & (e3 = op(e1, e1)) & (e3 = op(e0, e0))) | ((e2 = op(e3, e3)) & (e2 = op(e2, e2)) & (e2 = op(e1, e1)) & (e2 = op(e0, e0))) | ((e1 = op(e3, e3)) & (e1 = op(e2, e2)) & (e1 = op(e1, e1)) & (e1 = op(e0, e0))) | ((e0 = op(e3, e3)) & (e0 = op(e2, e2)) & (e0 = op(e1, e1)) & (e0 = op(e0, e0))))), inference(negated_conjecture, [], [f4])).
fof(f4, plain, ~ (((e3 = op(e3, e3)) | (e3 = op(e2, e3)) | (e3 = op(e1, e3)) | (e3 = op(e0, e3))) & ((e3 = op(e3, e3)) | (e3 = op(e3, e2)) | (e3 = op(e3, e1)) | (e3 = op(e3, e0))) & ((e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(e0, e3))) & ((e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, e0))) & ((e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(e0, e3))) & ((e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, e0))) & ((e0 = op(e3, e3)) | (e0 = op(e2, e3)) | (e0 = op(e1, e3)) | (e0 = op(e0, e3))) & ((e0 = op(e3, e3)) | (e0 = op(e3, e2)) | (e0 = op(e3, e1)) | (e0 = op(e3, e0))) & ((e3 = op(e3, e2)) | (e3 = op(e2, e2)) | (e3 = op(e1, e2)) | (e3 = op(e0, e2))) & ((e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, e0))) & ((e2 = op(e3, e2)) | (e2 = op(e2, e2)) | (e2 = op(e1, e2)) | (e2 = op(e0, e2))) & ((e2 = op(e2, e3)) | (e2 = op(e2, e2)) | (e2 = op(e2, e1)) | (e2 = op(e2, e0))) & ((e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(e0, e2))) & ((e1 = op(e2, e3)) | (e1 = op(e2, e2)) | (e1 = op(e2, e1)) | (e1 = op(e2, e0))) & ((e0 = op(e3, e2)) | (e0 = op(e2, e2)) | (e0 = op(e1, e2)) | (e0 = op(e0, e2))) & ((e0 = op(e2, e3)) | (e0 = op(e2, e2)) | (e0 = op(e2, e1)) | (e0 = op(e2, e0))) & ((e3 = op(e3, e1)) | (e3 = op(e2, e1)) | (e3 = op(e1, e1)) | (e3 = op(e0, e1))) & ((e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, e0))) & ((e2 = op(e3, e1)) | (e2 = op(e2, e1)) | (e2 = op(e1, e1)) | (e2 = op(e0, e1))) & ((e2 = op(e1, e3)) | (e2 = op(e1, e2)) | (e2 = op(e1, e1)) | (e2 = op(e1, e0))) & ((e1 = op(e3, e1)) | (e1 = op(e2, e1)) | (e1 = op(e1, e1)) | (e1 = op(e0, e1))) & ((e1 = op(e1, e3)) | (e1 = op(e1, e2)) | (e1 = op(e1, e1)) | (e1 = op(e1, e0))) & ((e0 = op(e3, e1)) | (e0 = op(e2, e1)) | (e0 = op(e1, e1)) | (e0 = op(e0, e1))) & ((e0 = op(e1, e3)) | (e0 = op(e1, e2)) | (e0 = op(e1, e1)) | (e0 = op(e1, e0))) & ((e3 = op(e3, e0)) | (e3 = op(e2, e0)) | (e3 = op(e1, e0)) | (e3 = op(e0, e0))) & ((e3 = op(e0, e3)) | (e3 = op(e0, e2)) | (e3 = op(e0, e1)) | (e3 = op(e0, e0))) & ((e2 = op(e3, e0)) | (e2 = op(e2, e0)) | (e2 = op(e1, e0)) | (e2 = op(e0, e0))) & ((e2 = op(e0, e3)) | (e2 = op(e0, e2)) | (e2 = op(e0, e1)) | (e2 = op(e0, e0))) & ((e1 = op(e3, e0)) | (e1 = op(e2, e0)) | (e1 = op(e1, e0)) | (e1 = op(e0, e0))) & ((e1 = op(e0, e3)) | (e1 = op(e0, e2)) | (e1 = op(e0, e1)) | (e1 = op(e0, e0))) & ((e0 = op(e3, e0)) | (e0 = op(e2, e0)) | (e0 = op(e1, e0)) | (e0 = op(e0, e0))) & ((e0 = op(e0, e3)) | (e0 = op(e0, e2)) | (e0 = op(e0, e1)) | (e0 = op(e0, e0))) & ((e3 = unit) | (e2 = unit) | (e1 = unit) | (e0 = unit)) & (e3 = op(e3, unit)) & (e3 = op(unit, e3)) & (e2 = op(e2, unit)) & (e2 = op(unit, e2)) & (e1 = op(e1, unit)) & (e1 = op(unit, e1)) & (e0 = op(e0, unit)) & (e0 = op(unit, e0)) & ((e3 = op(e3, e3)) | (e2 = op(e3, e3)) | (e1 = op(e3, e3)) | (e0 = op(e3, e3))) & ((e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (e0 = op(e3, e2))) & ((e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (e0 = op(e3, e1))) & ((e3 = op(e3, e0)) | (e2 = op(e3, e0)) | (e1 = op(e3, e0)) | (e0 = op(e3, e0))) & ((e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))) & ((e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (e0 = op(e2, e2))) & ((e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))) & ((e3 = op(e2, e0)) | (e2 = op(e2, e0)) | (e1 = op(e2, e0)) | (e0 = op(e2, e0))) & ((e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (e0 = op(e1, e3))) & ((e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (e0 = op(e1, e2))) & ((e3 = op(e1, e1)) | (e2 = op(e1, e1)) | (e1 = op(e1, e1)) | (e0 = op(e1, e1))) & ((e3 = op(e1, e0)) | (e2 = op(e1, e0)) | (e1 = op(e1, e0)) | (e0 = op(e1, e0))) & ((e3 = op(e0, e3)) | (e2 = op(e0, e3)) | (e1 = op(e0, e3)) | (e0 = op(e0, e3))) & ((e3 = op(e0, e2)) | (e2 = op(e0, e2)) | (e1 = op(e0, e2)) | (e0 = op(e0, e2))) & ((e3 = op(e0, e1)) | (e2 = op(e0, e1)) | (e1 = op(e0, e1)) | (e0 = op(e0, e1))) & ((e3 = op(e0, e0)) | (e2 = op(e0, e0)) | (e1 = op(e0, e0)) | (e0 = op(e0, e0))) & ~ (((e3 = op(e3, e3)) & (e3 = op(e2, e2)) & (e3 = op(e1, e1)) & (e3 = op(e0, e0))) | ((e2 = op(e3, e3)) & (e2 = op(e2, e2)) & (e2 = op(e1, e1)) & (e2 = op(e0, e0))) | ((e1 = op(e3, e3)) & (e1 = op(e2, e2)) & (e1 = op(e1, e1)) & (e1 = op(e0, e0))) | ((e0 = op(e3, e3)) & (e0 = op(e2, e2)) & (e0 = op(e1, e1)) & (e0 = op(e0, e0))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG044+1.p', co1)).