fof(f9990, plain, $false, inference(avatar_sat_refutation, [], [f2270, f2459, f2463, f2464, f2466, f2467, f2478, f2479, f2488, f2490, f2491, f2493, f2866, f3055, f3060, f3072, f3075, f3083, f3084, f3085, f3086, f3087, f3089, f3144, f3288, f3442, f3501, f3536, f3601, f3720, f3762, f3813, f3867, f3887, f3900, f3923, f3964, f3976, f4015, f4174, f4313, f4377, f4407, f4427, f4442, f4591, f4633, f4689, f4738, f4760, f4771, f4778, f4783, f4835, f4847, f4848, f4850, f4851, f4853, f5128, f5148, f5168, f5192, f5883, f6346, f6347, f6348, f6351, f6380, f6434, f6473, f6491, f6508, f6518, f6531, f6570, f6676, f6678, f6697, f6725, f6741, f6780, f6782, f6783, f6814, f6851, f6854, f6855, f6872, f6916, f7044, f7050, f7059, f7151, f7173, f7174, f7176, f7177, f7190, f7191, f7194, f7215, f7233, f7277, f7305, f7344, f7373, f7444, f7452, f7490, f7533, f7537, f7570, f7587, f7607, f7640, f7708, f7774, f7798, f7809, f7826, f7832, f7861, f7879, f7894, f7897, f7918, f7922, f7965, f8013, f8046, f8089, f8137, f8151, f8254, f8307, f8323, f8343, f8356, f8370, f8386, f8388, f8397, f8421, f8436, f8461, f8532, f8575, f8594, f8596, f8646, f8692, f8720, f8781, f8800, f8858, f8859, f8860, f8861, f8862, f8863, f8884, f8917, f8943, f8944, f8945, f8946, f8947, f8948, f8960, f8961, f8981, f8982, f9018, f9032, f9067, f9098, f9102, f9133, f9156, f9198, f9231, f9268, f9295, f9342, f9352, f9397, f9411, f9443, f9471, f9499, f9530, f9563, f9599, f9637, f9670, f9703, f9727, f9749, f9776, f9809, f9842, f9870, f9893, f9967, f9978])).
fof(f9978, plain, (~ spl280_221 | ~ spl280_224), inference(avatar_contradiction_clause, [], [f9977])).
fof(f9977, plain, ($false | (~ spl280_221 | ~ spl280_224)), inference(subsumption_resolution, [], [f9976, f971])).
fof(f971, plain, ~ (e20 = e23), inference(cnf_transformation, [], [f10])).
fof(f10, plain, (~ (e23 = e24) & ~ (e22 = e24) & ~ (e22 = e23) & ~ (e21 = e24) & ~ (e21 = e23) & ~ (e21 = e22) & ~ (e20 = e24) & ~ (e20 = e23) & ~ (e20 = e22) & ~ (e20 = e21)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG051+1.p', ax10)).
fof(f9976, plain, ((e20 = e23) | (~ spl280_221 | ~ spl280_224)), inference(forward_demodulation, [], [f2903, f2891])).
fof(f2891, plain, ((e20 = op2(e21, e21)) | ~ spl280_221), inference(avatar_component_clause, [], [f2889])).
fof(f2889, plain, (spl280_221 <=> (e20 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl280_221])])).
fof(f2903, plain, ((e23 = op2(e21, e21)) | ~ spl280_224), inference(avatar_component_clause, [], [f2901])).
fof(f2901, plain, (spl280_224 <=> (e23 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl280_224])])).
fof(f9967, plain, (~ spl280_528 | ~ spl280_3 | spl280_667), inference(avatar_split_clause, [], [f9966, f5784, f1923, f4953])).
fof(f4953, plain, (spl280_528 <=> (e22 = h5(e13))), introduced(avatar_definition, [new_symbols(naming, [spl280_528])])).
fof(f1923, plain, (spl280_3 <=> (e12 = op1(e14, e14))), introduced(avatar_definition, [new_symbols(naming, [spl280_3])])).
fof(f5784, plain, (spl280_667 <=> (h5(e13) = h3(op1(e14, e14)))), introduced(avatar_definition, [new_symbols(naming, [spl280_667])])).
fof(f9966, plain, (~ (e22 = h5(e13)) | (~ spl280_3 | spl280_667)), inference(forward_demodulation, [], [f9965, f1774])).
fof(f1774, plain, (e22 = h3(e12)), inference(cnf_transformation, [], [f18])).
fof(f18, plain, ((op2(op2(e22, e22), op2(e22, e22)) = h3(e14)) & (op2(e22, e22) = h3(e13)) & (op2(op2(e22, e22), e22) = h3(e11)) & (op2(op2(op2(e22, e22), e22), op2(op2(e22, e22), e22)) = h3(e10)) & (e22 = h3(e12))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG051+1.p', ax18)).
fof(f9965, plain, (~ (h3(e12) = h5(e13)) | (~ spl280_3 | spl280_667)), inference(forward_demodulation, [], [f5786, f1925])).
fof(f1925, plain, ((e12 = op1(e14, e14)) | ~ spl280_3), inference(avatar_component_clause, [], [f1923])).
fof(f5786, plain, (~ (h5(e13) = h3(op1(e14, e14))) | spl280_667), inference(avatar_component_clause, [], [f5784])).
fof(f9893, plain, (~ spl280_210 | ~ spl280_235), inference(avatar_split_clause, [], [f9892, f2947, f2842])).
fof(f2842, plain, (spl280_210 <=> (e24 = op2(e21, e24))), introduced(avatar_definition, [new_symbols(naming, [spl280_210])])).
fof(f2947, plain, (spl280_235 <=> (e24 = op2(e20, e24))), introduced(avatar_definition, [new_symbols(naming, [spl280_235])])).
fof(f9892, plain, (~ (e24 = op2(e21, e24)) | ~ spl280_235), inference(forward_demodulation, [], [f899, f2949])).
fof(f2949, plain, ((e24 = op2(e20, e24)) | ~ spl280_235), inference(avatar_component_clause, [], [f2947])).
fof(f899, plain, ~ (op2(e20, e24) = op2(e21, e24)), inference(cnf_transformation, [], [f8])).
fof(f8, plain, (~ (op2(e24, e23) = op2(e24, e24)) & ~ (op2(e24, e22) = op2(e24, e24)) & ~ (op2(e24, e22) = op2(e24, e23)) & ~ (op2(e24, e21) = op2(e24, e24)) & ~ (op2(e24, e21) = op2(e24, e23)) & ~ (op2(e24, e21) = op2(e24, e22)) & ~ (op2(e24, e20) = op2(e24, e24)) & ~ (op2(e24, e20) = op2(e24, e23)) & ~ (op2(e24, e20) = op2(e24, e22)) & ~ (op2(e24, e20) = op2(e24, e21)) & ~ (op2(e23, e23) = op2(e23, e24)) & ~ (op2(e23, e22) = op2(e23, e24)) & ~ (op2(e23, e22) = op2(e23, e23)) & ~ (op2(e23, e21) = op2(e23, e24)) & ~ (op2(e23, e21) = op2(e23, e23)) & ~ (op2(e23, e21) = op2(e23, e22)) & ~ (op2(e23, e20) = op2(e23, e24)) & ~ (op2(e23, e20) = op2(e23, e23)) & ~ (op2(e23, e20) = op2(e23, e22)) & ~ (op2(e23, e20) = op2(e23, e21)) & ~ (op2(e22, e23) = op2(e22, e24)) & ~ (op2(e22, e22) = op2(e22, e24)) & ~ (op2(e22, e22) = op2(e22, e23)) & ~ (op2(e22, e21) = op2(e22, e24)) & ~ (op2(e22, e21) = op2(e22, e23)) & ~ (op2(e22, e21) = op2(e22, e22)) & ~ (op2(e22, e20) = op2(e22, e24)) & ~ (op2(e22, e20) = op2(e22, e23)) & ~ (op2(e22, e20) = op2(e22, e22)) & ~ (op2(e22, e20) = op2(e22, e21)) & ~ (op2(e21, e23) = op2(e21, e24)) & ~ (op2(e21, e22) = op2(e21, e24)) & ~ (op2(e21, e22) = op2(e21, e23)) & ~ (op2(e21, e21) = op2(e21, e24)) & ~ (op2(e21, e21) = op2(e21, e23)) & ~ (op2(e21, e21) = op2(e21, e22)) & ~ (op2(e21, e20) = op2(e21, e24)) & ~ (op2(e21, e20) = op2(e21, e23)) & ~ (op2(e21, e20) = op2(e21, e22)) & ~ (op2(e21, e20) = op2(e21, e21)) & ~ (op2(e20, e23) = op2(e20, e24)) & ~ (op2(e20, e22) = op2(e20, e24)) & ~ (op2(e20, e22) = op2(e20, e23)) & ~ (op2(e20, e21) = op2(e20, e24)) & ~ (op2(e20, e21) = op2(e20, e23)) & ~ (op2(e20, e21) = op2(e20, e22)) & ~ (op2(e20, e20) = op2(e20, e24)) & ~ (op2(e20, e20) = op2(e20, e23)) & ~ (op2(e20, e20) = op2(e20, e22)) & ~ (op2(e20, e20) = op2(e20, e21)) & ~ (op2(e23, e24) = op2(e24, e24)) & ~ (op2(e22, e24) = op2(e24, e24)) & ~ (op2(e22, e24) = op2(e23, e24)) & ~ (op2(e21, e24) = op2(e24, e24)) & ~ (op2(e21, e24) = op2(e23, e24)) & ~ (op2(e21, e24) = op2(e22, e24)) & ~ (op2(e20, e24) = op2(e24, e24)) & ~ (op2(e20, e24) = op2(e23, e24)) & ~ (op2(e20, e24) = op2(e22, e24)) & ~ (op2(e20, e24) = op2(e21, e24)) & ~ (op2(e23, e23) = op2(e24, e23)) & ~ (op2(e22, e23) = op2(e24, e23)) & ~ (op2(e22, e23) = op2(e23, e23)) & ~ (op2(e21, e23) = op2(e24, e23)) & ~ (op2(e21, e23) = op2(e23, e23)) & ~ (op2(e21, e23) = op2(e22, e23)) & ~ (op2(e20, e23) = op2(e24, e23)) & ~ (op2(e20, e23) = op2(e23, e23)) & ~ (op2(e20, e23) = op2(e22, e23)) & ~ (op2(e20, e23) = op2(e21, e23)) & ~ (op2(e23, e22) = op2(e24, e22)) & ~ (op2(e22, e22) = op2(e24, e22)) & ~ (op2(e22, e22) = op2(e23, e22)) & ~ (op2(e21, e22) = op2(e24, e22)) & ~ (op2(e21, e22) = op2(e23, e22)) & ~ (op2(e21, e22) = op2(e22, e22)) & ~ (op2(e20, e22) = op2(e24, e22)) & ~ (op2(e20, e22) = op2(e23, e22)) & ~ (op2(e20, e22) = op2(e22, e22)) & ~ (op2(e20, e22) = op2(e21, e22)) & ~ (op2(e23, e21) = op2(e24, e21)) & ~ (op2(e22, e21) = op2(e24, e21)) & ~ (op2(e22, e21) = op2(e23, e21)) & ~ (op2(e21, e21) = op2(e24, e21)) & ~ (op2(e21, e21) = op2(e23, e21)) & ~ (op2(e21, e21) = op2(e22, e21)) & ~ (op2(e20, e21) = op2(e24, e21)) & ~ (op2(e20, e21) = op2(e23, e21)) & ~ (op2(e20, e21) = op2(e22, e21)) & ~ (op2(e20, e21) = op2(e21, e21)) & ~ (op2(e23, e20) = op2(e24, e20)) & ~ (op2(e22, e20) = op2(e24, e20)) & ~ (op2(e22, e20) = op2(e23, e20)) & ~ (op2(e21, e20) = op2(e24, e20)) & ~ (op2(e21, e20) = op2(e23, e20)) & ~ (op2(e21, e20) = op2(e22, e20)) & ~ (op2(e20, e20) = op2(e24, e20)) & ~ (op2(e20, e20) = op2(e23, e20)) & ~ (op2(e20, e20) = op2(e22, e20)) & ~ (op2(e20, e20) = op2(e21, e20))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG051+1.p', ax8)).
fof(f9870, plain, (~ spl280_105 | ~ spl280_235 | ~ spl280_572 | spl280_687), inference(avatar_contradiction_clause, [], [f9869])).
fof(f9869, plain, ($false | (~ spl280_105 | ~ spl280_235 | ~ spl280_572 | spl280_687)), inference(subsumption_resolution, [], [f9868, f2949])).
fof(f9868, plain, (~ (e24 = op2(e20, e24)) | (~ spl280_105 | ~ spl280_572 | spl280_687)), inference(forward_demodulation, [], [f9867, f4878])).
fof(f4878, plain, (e24 = h3(e14)), inference(forward_demodulation, [], [f1778, f1763])).
fof(f1763, plain, (e24 = op2(op2(e22, e22), op2(e22, e22))), inference(cnf_transformation, [], [f15])).
fof(f15, plain, ((e24 = op2(op2(e22, e22), op2(e22, e22))) & (e23 = op2(e22, e22)) & (e21 = op2(op2(e22, e22), e22)) & (e20 = op2(op2(op2(e22, e22), e22), op2(op2(e22, e22), e22)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG051+1.p', ax15)).
fof(f1778, plain, (op2(op2(e22, e22), op2(e22, e22)) = h3(e14)), inference(cnf_transformation, [], [f18])).
fof(f9867, plain, (~ (op2(e20, e24) = h3(e14)) | (~ spl280_105 | ~ spl280_572 | spl280_687)), inference(forward_demodulation, [], [f9866, f2353])).
fof(f2353, plain, ((e14 = op1(e10, e14)) | ~ spl280_105), inference(avatar_component_clause, [], [f2351])).
fof(f2351, plain, (spl280_105 <=> (e14 = op1(e10, e14))), introduced(avatar_definition, [new_symbols(naming, [spl280_105])])).
fof(f9866, plain, (~ (op2(e20, e24) = h3(op1(e10, e14))) | (~ spl280_572 | spl280_687)), inference(forward_demodulation, [], [f5866, f5190])).
fof(f5190, plain, ((e20 = h2(e13)) | ~ spl280_572), inference(avatar_component_clause, [], [f5189])).
fof(f5189, plain, (spl280_572 <=> (e20 = h2(e13))), introduced(avatar_definition, [new_symbols(naming, [spl280_572])])).
fof(f5866, plain, (~ (h3(op1(e10, e14)) = op2(h2(e13), e24)) | spl280_687), inference(avatar_component_clause, [], [f5864])).
fof(f5864, plain, (spl280_687 <=> (h3(op1(e10, e14)) = op2(h2(e13), e24))), introduced(avatar_definition, [new_symbols(naming, [spl280_687])])).
fof(f9842, plain, (~ spl280_97 | ~ spl280_227 | ~ spl280_572 | spl280_686), inference(avatar_contradiction_clause, [], [f9841])).
fof(f9841, plain, ($false | (~ spl280_97 | ~ spl280_227 | ~ spl280_572 | spl280_686)), inference(subsumption_resolution, [], [f9840, f2916])).
fof(f2916, plain, ((e21 = op2(e21, e20)) | ~ spl280_227), inference(avatar_component_clause, [], [f2914])).
fof(f2914, plain, (spl280_227 <=> (e21 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl280_227])])).
fof(f9840, plain, (~ (e21 = op2(e21, e20)) | (~ spl280_97 | ~ spl280_572 | spl280_686)), inference(forward_demodulation, [], [f9839, f4890])).
fof(f4890, plain, (e21 = h3(e11)), inference(forward_demodulation, [], [f4889, f4887])).
fof(f4887, plain, (e21 = op2(h3(e13), e22)), inference(backward_demodulation, [], [f1761, f1777])).
fof(f1777, plain, (op2(e22, e22) = h3(e13)), inference(cnf_transformation, [], [f18])).
fof(f1761, plain, (e21 = op2(op2(e22, e22), e22)), inference(cnf_transformation, [], [f15])).
fof(f4889, plain, (h3(e11) = op2(h3(e13), e22)), inference(forward_demodulation, [], [f1776, f1777])).
fof(f1776, plain, (op2(op2(e22, e22), e22) = h3(e11)), inference(cnf_transformation, [], [f18])).
fof(f9839, plain, (~ (op2(e21, e20) = h3(e11)) | (~ spl280_97 | ~ spl280_572 | spl280_686)), inference(forward_demodulation, [], [f9838, f2320])).
fof(f2320, plain, ((e11 = op1(e11, e10)) | ~ spl280_97), inference(avatar_component_clause, [], [f2318])).
fof(f2318, plain, (spl280_97 <=> (e11 = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl280_97])])).
fof(f9838, plain, (~ (op2(e21, e20) = h3(op1(e11, e10))) | (~ spl280_572 | spl280_686)), inference(forward_demodulation, [], [f5862, f5190])).
fof(f5862, plain, (~ (h3(op1(e11, e10)) = op2(e21, h2(e13))) | spl280_686), inference(avatar_component_clause, [], [f5860])).
fof(f5860, plain, (spl280_686 <=> (h3(op1(e11, e10)) = op2(e21, h2(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl280_686])])).
fof(f9809, plain, (~ spl280_90 | ~ spl280_220 | spl280_684), inference(avatar_contradiction_clause, [], [f9808])).
fof(f9808, plain, ($false | (~ spl280_90 | ~ spl280_220 | spl280_684)), inference(subsumption_resolution, [], [f9807, f2886])).
fof(f2886, plain, ((e24 = op2(e21, e22)) | ~ spl280_220), inference(avatar_component_clause, [], [f2884])).
fof(f2884, plain, (spl280_220 <=> (e24 = op2(e21, e22))), introduced(avatar_definition, [new_symbols(naming, [spl280_220])])).
fof(f9807, plain, (~ (e24 = op2(e21, e22)) | (~ spl280_90 | spl280_684)), inference(forward_demodulation, [], [f9806, f4878])).
fof(f9806, plain, (~ (op2(e21, e22) = h3(e14)) | (~ spl280_90 | spl280_684)), inference(forward_demodulation, [], [f5854, f2290])).
fof(f2290, plain, ((e14 = op1(e11, e12)) | ~ spl280_90), inference(avatar_component_clause, [], [f2288])).
fof(f2288, plain, (spl280_90 <=> (e14 = op1(e11, e12))), introduced(avatar_definition, [new_symbols(naming, [spl280_90])])).
fof(f5854, plain, (~ (op2(e21, e22) = h3(op1(e11, e12))) | spl280_684), inference(avatar_component_clause, [], [f5852])).
fof(f5852, plain, (spl280_684 <=> (op2(e21, e22) = h3(op1(e11, e12)))), introduced(avatar_definition, [new_symbols(naming, [spl280_684])])).
fof(f9776, plain, (~ spl280_73 | ~ spl280_203 | ~ spl280_572 | spl280_681), inference(avatar_contradiction_clause, [], [f9775])).
fof(f9775, plain, ($false | (~ spl280_73 | ~ spl280_203 | ~ spl280_572 | spl280_681)), inference(subsumption_resolution, [], [f9774, f2815])).
fof(f2815, plain, ((e22 = op2(e22, e20)) | ~ spl280_203), inference(avatar_component_clause, [], [f2813])).
fof(f2813, plain, (spl280_203 <=> (e22 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl280_203])])).
fof(f9774, plain, (~ (e22 = op2(e22, e20)) | (~ spl280_73 | ~ spl280_572 | spl280_681)), inference(forward_demodulation, [], [f9773, f1774])).
fof(f9773, plain, (~ (op2(e22, e20) = h3(e12)) | (~ spl280_73 | ~ spl280_572 | spl280_681)), inference(forward_demodulation, [], [f9772, f2219])).
fof(f2219, plain, ((e12 = op1(e12, e10)) | ~ spl280_73), inference(avatar_component_clause, [], [f2217])).
fof(f2217, plain, (spl280_73 <=> (e12 = op1(e12, e10))), introduced(avatar_definition, [new_symbols(naming, [spl280_73])])).
fof(f9772, plain, (~ (op2(e22, e20) = h3(op1(e12, e10))) | (~ spl280_572 | spl280_681)), inference(forward_demodulation, [], [f5842, f5190])).
fof(f5842, plain, (~ (h3(op1(e12, e10)) = op2(e22, h2(e13))) | spl280_681), inference(avatar_component_clause, [], [f5840])).
fof(f5840, plain, (spl280_681 <=> (h3(op1(e12, e10)) = op2(e22, h2(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl280_681])])).
fof(f9749, plain, (~ spl280_70 | ~ spl280_200 | spl280_680), inference(avatar_contradiction_clause, [], [f9748])).
fof(f9748, plain, ($false | (~ spl280_70 | ~ spl280_200 | spl280_680)), inference(subsumption_resolution, [], [f9747, f2802])).
fof(f2802, plain, ((e24 = op2(e22, e21)) | ~ spl280_200), inference(avatar_component_clause, [], [f2800])).
fof(f2800, plain, (spl280_200 <=> (e24 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl280_200])])).
fof(f9747, plain, (~ (e24 = op2(e22, e21)) | (~ spl280_70 | spl280_680)), inference(forward_demodulation, [], [f9746, f4878])).
fof(f9746, plain, (~ (op2(e22, e21) = h3(e14)) | (~ spl280_70 | spl280_680)), inference(forward_demodulation, [], [f5838, f2206])).
fof(f2206, plain, ((e14 = op1(e12, e11)) | ~ spl280_70), inference(avatar_component_clause, [], [f2204])).
fof(f2204, plain, (spl280_70 <=> (e14 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl280_70])])).
fof(f5838, plain, (~ (op2(e22, e21) = h3(op1(e12, e11))) | spl280_680), inference(avatar_component_clause, [], [f5836])).
fof(f5836, plain, (spl280_680 <=> (op2(e22, e21) = h3(op1(e12, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl280_680])])).
fof(f9727, plain, (~ spl280_64 | spl280_679), inference(avatar_contradiction_clause, [], [f9726])).
fof(f9726, plain, ($false | (~ spl280_64 | spl280_679)), inference(trivial_inequality_removal, [], [f9725])).
fof(f9725, plain, (~ (h3(e13) = h3(e13)) | (~ spl280_64 | spl280_679)), inference(forward_demodulation, [], [f5834, f2181])).
fof(f2181, plain, ((e13 = op1(e12, e12)) | ~ spl280_64), inference(avatar_component_clause, [], [f2179])).
fof(f2179, plain, (spl280_64 <=> (e13 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl280_64])])).
fof(f5834, plain, (~ (h3(e13) = h3(op1(e12, e12))) | spl280_679), inference(avatar_component_clause, [], [f5832])).
fof(f5832, plain, (spl280_679 <=> (h3(e13) = h3(op1(e12, e12)))), introduced(avatar_definition, [new_symbols(naming, [spl280_679])])).
fof(f9703, plain, (~ spl280_52 | ~ spl280_182 | spl280_677), inference(avatar_contradiction_clause, [], [f9702])).
fof(f9702, plain, ($false | (~ spl280_52 | ~ spl280_182 | spl280_677)), inference(subsumption_resolution, [], [f9701, f2727])).
fof(f2727, plain, ((e21 = op2(e22, e24)) | ~ spl280_182), inference(avatar_component_clause, [], [f2725])).
fof(f2725, plain, (spl280_182 <=> (e21 = op2(e22, e24))), introduced(avatar_definition, [new_symbols(naming, [spl280_182])])).
fof(f9701, plain, (~ (e21 = op2(e22, e24)) | (~ spl280_52 | spl280_677)), inference(forward_demodulation, [], [f9700, f4890])).
fof(f9700, plain, (~ (op2(e22, e24) = h3(e11)) | (~ spl280_52 | spl280_677)), inference(forward_demodulation, [], [f5826, f2131])).
fof(f2131, plain, ((e11 = op1(e12, e14)) | ~ spl280_52), inference(avatar_component_clause, [], [f2129])).
fof(f2129, plain, (spl280_52 <=> (e11 = op1(e12, e14))), introduced(avatar_definition, [new_symbols(naming, [spl280_52])])).
fof(f5826, plain, (~ (op2(e22, e24) = h3(op1(e12, e14))) | spl280_677), inference(avatar_component_clause, [], [f5824])).
fof(f5824, plain, (spl280_677 <=> (op2(e22, e24) = h3(op1(e12, e14)))), introduced(avatar_definition, [new_symbols(naming, [spl280_677])])).
fof(f9670, plain, (~ spl280_117 | ~ spl280_583 | spl280_690), inference(avatar_contradiction_clause, [], [f9669])).
fof(f9669, plain, ($false | (~ spl280_117 | ~ spl280_583 | spl280_690)), inference(subsumption_resolution, [], [f9668, f5249])).
fof(f5249, plain, ((e21 = h2(e11)) | ~ spl280_583), inference(avatar_component_clause, [], [f5248])).
fof(f5248, plain, (spl280_583 <=> (e21 = h2(e11))), introduced(avatar_definition, [new_symbols(naming, [spl280_583])])).
fof(f9668, plain, (~ (e21 = h2(e11)) | (~ spl280_117 | spl280_690)), inference(forward_demodulation, [], [f9667, f4890])).
fof(f9667, plain, (~ (h2(e11) = h3(e11)) | (~ spl280_117 | spl280_690)), inference(forward_demodulation, [], [f5878, f2404])).
fof(f2404, plain, ((e11 = op1(e10, e11)) | ~ spl280_117), inference(avatar_component_clause, [], [f2402])).
fof(f2402, plain, (spl280_117 <=> (e11 = op1(e10, e11))), introduced(avatar_definition, [new_symbols(naming, [spl280_117])])).
fof(f5878, plain, (~ (h2(e11) = h3(op1(e10, e11))) | spl280_690), inference(avatar_component_clause, [], [f5876])).
fof(f5876, plain, (spl280_690 <=> (h2(e11) = h3(op1(e10, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl280_690])])).
fof(f9637, plain, (~ spl280_83 | ~ spl280_213 | ~ spl280_562 | spl280_683), inference(avatar_contradiction_clause, [], [f9636])).
fof(f9636, plain, ($false | (~ spl280_83 | ~ spl280_213 | ~ spl280_562 | spl280_683)), inference(subsumption_resolution, [], [f9635, f2857])).
fof(f2857, plain, ((e22 = op2(e21, e23)) | ~ spl280_213), inference(avatar_component_clause, [], [f2855])).
fof(f2855, plain, (spl280_213 <=> (e22 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl280_213])])).
fof(f9635, plain, (~ (e22 = op2(e21, e23)) | (~ spl280_83 | ~ spl280_562 | spl280_683)), inference(forward_demodulation, [], [f9634, f1774])).
fof(f9634, plain, (~ (op2(e21, e23) = h3(e12)) | (~ spl280_83 | ~ spl280_562 | spl280_683)), inference(forward_demodulation, [], [f9633, f2261])).
fof(f2261, plain, ((e12 = op1(e11, e13)) | ~ spl280_83), inference(avatar_component_clause, [], [f2259])).
fof(f2259, plain, (spl280_83 <=> (e12 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl280_83])])).
fof(f9633, plain, (~ (op2(e21, e23) = h3(op1(e11, e13))) | (~ spl280_562 | spl280_683)), inference(forward_demodulation, [], [f5850, f5126])).
fof(f5126, plain, ((e23 = h3(e13)) | ~ spl280_562), inference(avatar_component_clause, [], [f5125])).
fof(f5125, plain, (spl280_562 <=> (e23 = h3(e13))), introduced(avatar_definition, [new_symbols(naming, [spl280_562])])).
fof(f5850, plain, (~ (h3(op1(e11, e13)) = op2(e21, h3(e13))) | spl280_683), inference(avatar_component_clause, [], [f5848])).
fof(f5848, plain, (spl280_683 <=> (h3(op1(e11, e13)) = op2(e21, h3(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl280_683])])).
fof(f9599, plain, (~ spl280_56 | ~ spl280_186 | ~ spl280_562 | ~ spl280_572 | spl280_678), inference(avatar_contradiction_clause, [], [f9598])).
fof(f9598, plain, ($false | (~ spl280_56 | ~ spl280_186 | ~ spl280_562 | ~ spl280_572 | spl280_678)), inference(subsumption_resolution, [], [f9597, f2744])).
fof(f2744, plain, ((e20 = op2(e22, e23)) | ~ spl280_186), inference(avatar_component_clause, [], [f2742])).
fof(f2742, plain, (spl280_186 <=> (e20 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl280_186])])).
fof(f9597, plain, (~ (e20 = op2(e22, e23)) | (~ spl280_56 | ~ spl280_562 | ~ spl280_572 | spl280_678)), inference(forward_demodulation, [], [f9596, f7170])).
fof(f7170, plain, ((e20 = h3(e10)) | ~ spl280_572), inference(backward_demodulation, [], [f4893, f5190])).
fof(f4893, plain, (h2(e13) = h3(e10)), inference(forward_demodulation, [], [f4892, f1772])).
fof(f1772, plain, (op2(e21, e21) = h2(e13)), inference(cnf_transformation, [], [f17])).
fof(f17, plain, ((h2(e14) = op2(op2(e21, e21), op2(e21, e21))) & (op2(e21, e21) = h2(e13)) & (op2(op2(e21, e21), e21) = h2(e11)) & (h2(e10) = op2(op2(op2(e21, e21), e21), op2(op2(e21, e21), e21))) & (e21 = h2(e12))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG051+1.p', ax17)).
fof(f4892, plain, (op2(e21, e21) = h3(e10)), inference(forward_demodulation, [], [f4891, f4887])).
fof(f4891, plain, (h3(e10) = op2(op2(h3(e13), e22), op2(h3(e13), e22))), inference(forward_demodulation, [], [f1775, f1777])).
fof(f1775, plain, (op2(op2(op2(e22, e22), e22), op2(op2(e22, e22), e22)) = h3(e10)), inference(cnf_transformation, [], [f18])).
fof(f9596, plain, (~ (op2(e22, e23) = h3(e10)) | (~ spl280_56 | ~ spl280_562 | spl280_678)), inference(forward_demodulation, [], [f9595, f2148])).
fof(f2148, plain, ((e10 = op1(e12, e13)) | ~ spl280_56), inference(avatar_component_clause, [], [f2146])).
fof(f2146, plain, (spl280_56 <=> (e10 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl280_56])])).
fof(f9595, plain, (~ (op2(e22, e23) = h3(op1(e12, e13))) | (~ spl280_562 | spl280_678)), inference(forward_demodulation, [], [f5830, f5126])).
fof(f5830, plain, (~ (h3(op1(e12, e13)) = op2(e22, h3(e13))) | spl280_678), inference(avatar_component_clause, [], [f5828])).
fof(f5828, plain, (spl280_678 <=> (h3(op1(e12, e13)) = op2(e22, h3(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl280_678])])).
fof(f9563, plain, (~ spl280_49 | ~ spl280_179 | ~ spl280_562 | ~ spl280_572 | spl280_676), inference(avatar_contradiction_clause, [], [f9562])).
fof(f9562, plain, ($false | (~ spl280_49 | ~ spl280_179 | ~ spl280_562 | ~ spl280_572 | spl280_676)), inference(subsumption_resolution, [], [f9561, f2714])).
fof(f2714, plain, ((e23 = op2(e23, e20)) | ~ spl280_179), inference(avatar_component_clause, [], [f2712])).
fof(f2712, plain, (spl280_179 <=> (e23 = op2(e23, e20))), introduced(avatar_definition, [new_symbols(naming, [spl280_179])])).
fof(f9561, plain, (~ (e23 = op2(e23, e20)) | (~ spl280_49 | ~ spl280_562 | ~ spl280_572 | spl280_676)), inference(forward_demodulation, [], [f9560, f5126])).
fof(f9560, plain, (~ (op2(e23, e20) = h3(e13)) | (~ spl280_49 | ~ spl280_562 | ~ spl280_572 | spl280_676)), inference(forward_demodulation, [], [f9559, f2118])).
fof(f2118, plain, ((e13 = op1(e13, e10)) | ~ spl280_49), inference(avatar_component_clause, [], [f2116])).
fof(f2116, plain, (spl280_49 <=> (e13 = op1(e13, e10))), introduced(avatar_definition, [new_symbols(naming, [spl280_49])])).
fof(f9559, plain, (~ (op2(e23, e20) = h3(op1(e13, e10))) | (~ spl280_562 | ~ spl280_572 | spl280_676)), inference(forward_demodulation, [], [f9558, f5126])).
fof(f9558, plain, (~ (h3(op1(e13, e10)) = op2(h3(e13), e20)) | (~ spl280_572 | spl280_676)), inference(forward_demodulation, [], [f5822, f5190])).
fof(f5822, plain, (~ (h3(op1(e13, e10)) = op2(h3(e13), h2(e13))) | spl280_676), inference(avatar_component_clause, [], [f5820])).
fof(f5820, plain, (spl280_676 <=> (h3(op1(e13, e10)) = op2(h3(e13), h2(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl280_676])])).
fof(f9530, plain, (~ spl280_37 | spl280_674), inference(avatar_contradiction_clause, [], [f9529])).
fof(f9529, plain, ($false | (~ spl280_37 | spl280_674)), inference(subsumption_resolution, [], [f9528, f4890])).
fof(f9528, plain, (~ (e21 = h3(e11)) | (~ spl280_37 | spl280_674)), inference(forward_demodulation, [], [f5814, f2068])).
fof(f2068, plain, ((e11 = op1(e13, e12)) | ~ spl280_37), inference(avatar_component_clause, [], [f2066])).
fof(f2066, plain, (spl280_37 <=> (e11 = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl280_37])])).
fof(f5814, plain, (~ (e21 = h3(op1(e13, e12))) | spl280_674), inference(avatar_component_clause, [], [f5812])).
fof(f5812, plain, (spl280_674 <=> (e21 = h3(op1(e13, e12)))), introduced(avatar_definition, [new_symbols(naming, [spl280_674])])).
fof(f9499, plain, (~ spl280_35 | spl280_673), inference(avatar_contradiction_clause, [], [f9498])).
fof(f9498, plain, ($false | (~ spl280_35 | spl280_673)), inference(subsumption_resolution, [], [f9497, f4878])).
fof(f9497, plain, (~ (e24 = h3(e14)) | (~ spl280_35 | spl280_673)), inference(forward_demodulation, [], [f5810, f2059])).
fof(f2059, plain, ((e14 = op1(e13, e13)) | ~ spl280_35), inference(avatar_component_clause, [], [f2057])).
fof(f2057, plain, (spl280_35 <=> (e14 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl280_35])])).
fof(f5810, plain, (~ (e24 = h3(op1(e13, e13))) | spl280_673), inference(avatar_component_clause, [], [f5808])).
fof(f5808, plain, (spl280_673 <=> (e24 = h3(op1(e13, e13)))), introduced(avatar_definition, [new_symbols(naming, [spl280_673])])).
fof(f9471, plain, (~ spl280_25 | ~ spl280_155 | ~ spl280_572 | spl280_671), inference(avatar_contradiction_clause, [], [f9470])).
fof(f9470, plain, ($false | (~ spl280_25 | ~ spl280_155 | ~ spl280_572 | spl280_671)), inference(subsumption_resolution, [], [f9469, f2613])).
fof(f2613, plain, ((e24 = op2(e24, e20)) | ~ spl280_155), inference(avatar_component_clause, [], [f2611])).
fof(f2611, plain, (spl280_155 <=> (e24 = op2(e24, e20))), introduced(avatar_definition, [new_symbols(naming, [spl280_155])])).
fof(f9469, plain, (~ (e24 = op2(e24, e20)) | (~ spl280_25 | ~ spl280_572 | spl280_671)), inference(forward_demodulation, [], [f9468, f4878])).
fof(f9468, plain, (~ (op2(e24, e20) = h3(e14)) | (~ spl280_25 | ~ spl280_572 | spl280_671)), inference(forward_demodulation, [], [f9467, f2017])).
fof(f2017, plain, ((e14 = op1(e14, e10)) | ~ spl280_25), inference(avatar_component_clause, [], [f2015])).
fof(f2015, plain, (spl280_25 <=> (e14 = op1(e14, e10))), introduced(avatar_definition, [new_symbols(naming, [spl280_25])])).
fof(f9467, plain, (~ (op2(e24, e20) = h3(op1(e14, e10))) | (~ spl280_572 | spl280_671)), inference(forward_demodulation, [], [f5802, f5190])).
fof(f5802, plain, (~ (h3(op1(e14, e10)) = op2(e24, h2(e13))) | spl280_671), inference(avatar_component_clause, [], [f5800])).
fof(f5800, plain, (spl280_671 <=> (h3(op1(e14, e10)) = op2(e24, h2(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl280_671])])).
fof(f9443, plain, (~ spl280_19 | ~ spl280_149 | ~ spl280_562 | spl280_670), inference(avatar_contradiction_clause, [], [f9442])).
fof(f9442, plain, ($false | (~ spl280_19 | ~ spl280_149 | ~ spl280_562 | spl280_670)), inference(subsumption_resolution, [], [f9441, f2588])).
fof(f2588, plain, ((e23 = op2(e24, e21)) | ~ spl280_149), inference(avatar_component_clause, [], [f2586])).
fof(f2586, plain, (spl280_149 <=> (e23 = op2(e24, e21))), introduced(avatar_definition, [new_symbols(naming, [spl280_149])])).
fof(f9441, plain, (~ (e23 = op2(e24, e21)) | (~ spl280_19 | ~ spl280_562 | spl280_670)), inference(forward_demodulation, [], [f9440, f5126])).
fof(f9440, plain, (~ (op2(e24, e21) = h3(e13)) | (~ spl280_19 | spl280_670)), inference(forward_demodulation, [], [f5798, f1992])).
fof(f1992, plain, ((e13 = op1(e14, e11)) | ~ spl280_19), inference(avatar_component_clause, [], [f1990])).
fof(f1990, plain, (spl280_19 <=> (e13 = op1(e14, e11))), introduced(avatar_definition, [new_symbols(naming, [spl280_19])])).
fof(f5798, plain, (~ (op2(e24, e21) = h3(op1(e14, e11))) | spl280_670), inference(avatar_component_clause, [], [f5796])).
fof(f5796, plain, (spl280_670 <=> (op2(e24, e21) = h3(op1(e14, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl280_670])])).
fof(f9411, plain, (~ spl280_11 | ~ spl280_141 | ~ spl280_572 | spl280_669), inference(avatar_contradiction_clause, [], [f9410])).
fof(f9410, plain, ($false | (~ spl280_11 | ~ spl280_141 | ~ spl280_572 | spl280_669)), inference(subsumption_resolution, [], [f9409, f2555])).
fof(f2555, plain, ((e20 = op2(e24, e22)) | ~ spl280_141), inference(avatar_component_clause, [], [f2553])).
fof(f2553, plain, (spl280_141 <=> (e20 = op2(e24, e22))), introduced(avatar_definition, [new_symbols(naming, [spl280_141])])).
fof(f9409, plain, (~ (e20 = op2(e24, e22)) | (~ spl280_11 | ~ spl280_572 | spl280_669)), inference(forward_demodulation, [], [f9408, f7170])).
fof(f9408, plain, (~ (op2(e24, e22) = h3(e10)) | (~ spl280_11 | spl280_669)), inference(forward_demodulation, [], [f5794, f1959])).
fof(f1959, plain, ((e10 = op1(e14, e12)) | ~ spl280_11), inference(avatar_component_clause, [], [f1957])).
fof(f1957, plain, (spl280_11 <=> (e10 = op1(e14, e12))), introduced(avatar_definition, [new_symbols(naming, [spl280_11])])).
fof(f5794, plain, (~ (op2(e24, e22) = h3(op1(e14, e12))) | spl280_669), inference(avatar_component_clause, [], [f5792])).
fof(f5792, plain, (spl280_669 <=> (op2(e24, e22) = h3(op1(e14, e12)))), introduced(avatar_definition, [new_symbols(naming, [spl280_669])])).
fof(f9397, plain, (~ spl280_7 | ~ spl280_137 | ~ spl280_562 | spl280_668), inference(avatar_contradiction_clause, [], [f9396])).
fof(f9396, plain, ($false | (~ spl280_7 | ~ spl280_137 | ~ spl280_562 | spl280_668)), inference(subsumption_resolution, [], [f9392, f4890])).
fof(f9392, plain, (~ (e21 = h3(e11)) | (~ spl280_7 | ~ spl280_137 | ~ spl280_562 | spl280_668)), inference(backward_demodulation, [], [f9321, f1942])).
fof(f1942, plain, ((e11 = op1(e14, e13)) | ~ spl280_7), inference(avatar_component_clause, [], [f1940])).
fof(f1940, plain, (spl280_7 <=> (e11 = op1(e14, e13))), introduced(avatar_definition, [new_symbols(naming, [spl280_7])])).
fof(f9321, plain, (~ (e21 = h3(op1(e14, e13))) | (~ spl280_137 | ~ spl280_562 | spl280_668)), inference(forward_demodulation, [], [f9320, f2538])).
fof(f2538, plain, ((e21 = op2(e24, e23)) | ~ spl280_137), inference(avatar_component_clause, [], [f2536])).
fof(f2536, plain, (spl280_137 <=> (e21 = op2(e24, e23))), introduced(avatar_definition, [new_symbols(naming, [spl280_137])])).
fof(f9320, plain, (~ (op2(e24, e23) = h3(op1(e14, e13))) | (~ spl280_562 | spl280_668)), inference(forward_demodulation, [], [f5790, f5126])).
fof(f5790, plain, (~ (h3(op1(e14, e13)) = op2(e24, h3(e13))) | spl280_668), inference(avatar_component_clause, [], [f5788])).
fof(f5788, plain, (spl280_668 <=> (h3(op1(e14, e13)) = op2(e24, h3(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl280_668])])).
fof(f9352, plain, (~ spl280_11 | ~ spl280_12), inference(avatar_contradiction_clause, [], [f9351])).
fof(f9351, plain, ($false | (~ spl280_11 | ~ spl280_12)), inference(subsumption_resolution, [], [f9350, f959])).
fof(f959, plain, ~ (e10 = e11), inference(cnf_transformation, [], [f9])).
fof(f9, plain, (~ (e13 = e14) & ~ (e12 = e14) & ~ (e12 = e13) & ~ (e11 = e14) & ~ (e11 = e13) & ~ (e11 = e12) & ~ (e10 = e14) & ~ (e10 = e13) & ~ (e10 = e12) & ~ (e10 = e11)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG051+1.p', ax9)).
fof(f9350, plain, ((e10 = e11) | (~ spl280_11 | ~ spl280_12)), inference(forward_demodulation, [], [f1963, f1959])).
fof(f1963, plain, ((e11 = op1(e14, e12)) | ~ spl280_12), inference(avatar_component_clause, [], [f1961])).
fof(f1961, plain, (spl280_12 <=> (e11 = op1(e14, e12))), introduced(avatar_definition, [new_symbols(naming, [spl280_12])])).
fof(f9342, plain, (~ spl280_26 | ~ spl280_28), inference(avatar_contradiction_clause, [], [f9341])).
fof(f9341, plain, ($false | (~ spl280_26 | ~ spl280_28)), inference(subsumption_resolution, [], [f9340, f960])).
fof(f960, plain, ~ (e10 = e12), inference(cnf_transformation, [], [f9])).
fof(f9340, plain, ((e10 = e12) | (~ spl280_26 | ~ spl280_28)), inference(forward_demodulation, [], [f2030, f2022])).
fof(f2022, plain, ((e10 = op1(e13, e14)) | ~ spl280_26), inference(avatar_component_clause, [], [f2020])).
fof(f2020, plain, (spl280_26 <=> (e10 = op1(e13, e14))), introduced(avatar_definition, [new_symbols(naming, [spl280_26])])).
fof(f2030, plain, ((e12 = op1(e13, e14)) | ~ spl280_28), inference(avatar_component_clause, [], [f2028])).
fof(f2028, plain, (spl280_28 <=> (e12 = op1(e13, e14))), introduced(avatar_definition, [new_symbols(naming, [spl280_28])])).
fof(f9295, plain, (~ spl280_4 | ~ spl280_79), inference(avatar_split_clause, [], [f9294, f2242, f1927])).
fof(f1927, plain, (spl280_4 <=> (e13 = op1(e14, e14))), introduced(avatar_definition, [new_symbols(naming, [spl280_4])])).
fof(f2242, plain, (spl280_79 <=> (e13 = op1(e11, e14))), introduced(avatar_definition, [new_symbols(naming, [spl280_79])])).
fof(f9294, plain, (~ (e13 = op1(e14, e14)) | ~ spl280_79), inference(forward_demodulation, [], [f805, f2244])).
fof(f2244, plain, ((e13 = op1(e11, e14)) | ~ spl280_79), inference(avatar_component_clause, [], [f2242])).
fof(f805, plain, ~ (op1(e11, e14) = op1(e14, e14)), inference(cnf_transformation, [], [f7])).
fof(f7, plain, (~ (op1(e14, e13) = op1(e14, e14)) & ~ (op1(e14, e12) = op1(e14, e14)) & ~ (op1(e14, e12) = op1(e14, e13)) & ~ (op1(e14, e11) = op1(e14, e14)) & ~ (op1(e14, e11) = op1(e14, e13)) & ~ (op1(e14, e11) = op1(e14, e12)) & ~ (op1(e14, e10) = op1(e14, e14)) & ~ (op1(e14, e10) = op1(e14, e13)) & ~ (op1(e14, e10) = op1(e14, e12)) & ~ (op1(e14, e10) = op1(e14, e11)) & ~ (op1(e13, e13) = op1(e13, e14)) & ~ (op1(e13, e12) = op1(e13, e14)) & ~ (op1(e13, e12) = op1(e13, e13)) & ~ (op1(e13, e11) = op1(e13, e14)) & ~ (op1(e13, e11) = op1(e13, e13)) & ~ (op1(e13, e11) = op1(e13, e12)) & ~ (op1(e13, e10) = op1(e13, e14)) & ~ (op1(e13, e10) = op1(e13, e13)) & ~ (op1(e13, e10) = op1(e13, e12)) & ~ (op1(e13, e10) = op1(e13, e11)) & ~ (op1(e12, e13) = op1(e12, e14)) & ~ (op1(e12, e12) = op1(e12, e14)) & ~ (op1(e12, e12) = op1(e12, e13)) & ~ (op1(e12, e11) = op1(e12, e14)) & ~ (op1(e12, e11) = op1(e12, e13)) & ~ (op1(e12, e11) = op1(e12, e12)) & ~ (op1(e12, e10) = op1(e12, e14)) & ~ (op1(e12, e10) = op1(e12, e13)) & ~ (op1(e12, e10) = op1(e12, e12)) & ~ (op1(e12, e10) = op1(e12, e11)) & ~ (op1(e11, e13) = op1(e11, e14)) & ~ (op1(e11, e12) = op1(e11, e14)) & ~ (op1(e11, e12) = op1(e11, e13)) & ~ (op1(e11, e11) = op1(e11, e14)) & ~ (op1(e11, e11) = op1(e11, e13)) & ~ (op1(e11, e11) = op1(e11, e12)) & ~ (op1(e11, e10) = op1(e11, e14)) & ~ (op1(e11, e10) = op1(e11, e13)) & ~ (op1(e11, e10) = op1(e11, e12)) & ~ (op1(e11, e10) = op1(e11, e11)) & ~ (op1(e10, e13) = op1(e10, e14)) & ~ (op1(e10, e12) = op1(e10, e14)) & ~ (op1(e10, e12) = op1(e10, e13)) & ~ (op1(e10, e11) = op1(e10, e14)) & ~ (op1(e10, e11) = op1(e10, e13)) & ~ (op1(e10, e11) = op1(e10, e12)) & ~ (op1(e10, e10) = op1(e10, e14)) & ~ (op1(e10, e10) = op1(e10, e13)) & ~ (op1(e10, e10) = op1(e10, e12)) & ~ (op1(e10, e10) = op1(e10, e11)) & ~ (op1(e13, e14) = op1(e14, e14)) & ~ (op1(e12, e14) = op1(e14, e14)) & ~ (op1(e12, e14) = op1(e13, e14)) & ~ (op1(e11, e14) = op1(e14, e14)) & ~ (op1(e11, e14) = op1(e13, e14)) & ~ (op1(e11, e14) = op1(e12, e14)) & ~ (op1(e10, e14) = op1(e14, e14)) & ~ (op1(e10, e14) = op1(e13, e14)) & ~ (op1(e10, e14) = op1(e12, e14)) & ~ (op1(e10, e14) = op1(e11, e14)) & ~ (op1(e13, e13) = op1(e14, e13)) & ~ (op1(e12, e13) = op1(e14, e13)) & ~ (op1(e12, e13) = op1(e13, e13)) & ~ (op1(e11, e13) = op1(e14, e13)) & ~ (op1(e11, e13) = op1(e13, e13)) & ~ (op1(e11, e13) = op1(e12, e13)) & ~ (op1(e10, e13) = op1(e14, e13)) & ~ (op1(e10, e13) = op1(e13, e13)) & ~ (op1(e10, e13) = op1(e12, e13)) & ~ (op1(e10, e13) = op1(e11, e13)) & ~ (op1(e13, e12) = op1(e14, e12)) & ~ (op1(e12, e12) = op1(e14, e12)) & ~ (op1(e12, e12) = op1(e13, e12)) & ~ (op1(e11, e12) = op1(e14, e12)) & ~ (op1(e11, e12) = op1(e13, e12)) & ~ (op1(e11, e12) = op1(e12, e12)) & ~ (op1(e10, e12) = op1(e14, e12)) & ~ (op1(e10, e12) = op1(e13, e12)) & ~ (op1(e10, e12) = op1(e12, e12)) & ~ (op1(e10, e12) = op1(e11, e12)) & ~ (op1(e13, e11) = op1(e14, e11)) & ~ (op1(e12, e11) = op1(e14, e11)) & ~ (op1(e12, e11) = op1(e13, e11)) & ~ (op1(e11, e11) = op1(e14, e11)) & ~ (op1(e11, e11) = op1(e13, e11)) & ~ (op1(e11, e11) = op1(e12, e11)) & ~ (op1(e10, e11) = op1(e14, e11)) & ~ (op1(e10, e11) = op1(e13, e11)) & ~ (op1(e10, e11) = op1(e12, e11)) & ~ (op1(e10, e11) = op1(e11, e11)) & ~ (op1(e13, e10) = op1(e14, e10)) & ~ (op1(e12, e10) = op1(e14, e10)) & ~ (op1(e12, e10) = op1(e13, e10)) & ~ (op1(e11, e10) = op1(e14, e10)) & ~ (op1(e11, e10) = op1(e13, e10)) & ~ (op1(e11, e10) = op1(e12, e10)) & ~ (op1(e10, e10) = op1(e14, e10)) & ~ (op1(e10, e10) = op1(e13, e10)) & ~ (op1(e10, e10) = op1(e12, e10)) & ~ (op1(e10, e10) = op1(e11, e10))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG051+1.p', ax7)).
fof(f9268, plain, (~ spl280_113 | ~ spl280_243 | ~ spl280_572 | spl280_689), inference(avatar_contradiction_clause, [], [f9267])).
fof(f9267, plain, ($false | (~ spl280_113 | ~ spl280_243 | ~ spl280_572 | spl280_689)), inference(subsumption_resolution, [], [f9266, f2983])).
fof(f2983, plain, ((e22 = op2(e20, e22)) | ~ spl280_243), inference(avatar_component_clause, [], [f2981])).
fof(f2981, plain, (spl280_243 <=> (e22 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl280_243])])).
fof(f9266, plain, (~ (e22 = op2(e20, e22)) | (~ spl280_113 | ~ spl280_572 | spl280_689)), inference(forward_demodulation, [], [f9265, f1774])).
fof(f9265, plain, (~ (op2(e20, e22) = h3(e12)) | (~ spl280_113 | ~ spl280_572 | spl280_689)), inference(forward_demodulation, [], [f9264, f2387])).
fof(f2387, plain, ((e12 = op1(e10, e12)) | ~ spl280_113), inference(avatar_component_clause, [], [f2385])).
fof(f2385, plain, (spl280_113 <=> (e12 = op1(e10, e12))), introduced(avatar_definition, [new_symbols(naming, [spl280_113])])).
fof(f9264, plain, (~ (op2(e20, e22) = h3(op1(e10, e12))) | (~ spl280_572 | spl280_689)), inference(forward_demodulation, [], [f5874, f5190])).
fof(f5874, plain, (~ (h3(op1(e10, e12)) = op2(h2(e13), e22)) | spl280_689), inference(avatar_component_clause, [], [f5872])).
fof(f5872, plain, (spl280_689 <=> (h3(op1(e10, e12)) = op2(h2(e13), e22))), introduced(avatar_definition, [new_symbols(naming, [spl280_689])])).
fof(f9231, plain, (~ spl280_91 | ~ spl280_572 | spl280_685), inference(avatar_contradiction_clause, [], [f9230])).
fof(f9230, plain, ($false | (~ spl280_91 | ~ spl280_572 | spl280_685)), inference(subsumption_resolution, [], [f9229, f5190])).
fof(f9229, plain, (~ (e20 = h2(e13)) | (~ spl280_91 | ~ spl280_572 | spl280_685)), inference(forward_demodulation, [], [f9228, f7170])).
fof(f9228, plain, (~ (h2(e13) = h3(e10)) | (~ spl280_91 | spl280_685)), inference(forward_demodulation, [], [f5858, f2295])).
fof(f2295, plain, ((e10 = op1(e11, e11)) | ~ spl280_91), inference(avatar_component_clause, [], [f2293])).
fof(f2293, plain, (spl280_91 <=> (e10 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl280_91])])).
fof(f5858, plain, (~ (h2(e13) = h3(op1(e11, e11))) | spl280_685), inference(avatar_component_clause, [], [f5856])).
fof(f5856, plain, (spl280_685 <=> (h2(e13) = h3(op1(e11, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl280_685])])).
fof(f9198, plain, (~ spl280_79 | ~ spl280_209 | ~ spl280_562 | spl280_682), inference(avatar_contradiction_clause, [], [f9197])).
fof(f9197, plain, ($false | (~ spl280_79 | ~ spl280_209 | ~ spl280_562 | spl280_682)), inference(subsumption_resolution, [], [f9196, f2840])).
fof(f2840, plain, ((e23 = op2(e21, e24)) | ~ spl280_209), inference(avatar_component_clause, [], [f2838])).
fof(f2838, plain, (spl280_209 <=> (e23 = op2(e21, e24))), introduced(avatar_definition, [new_symbols(naming, [spl280_209])])).
fof(f9196, plain, (~ (e23 = op2(e21, e24)) | (~ spl280_79 | ~ spl280_562 | spl280_682)), inference(forward_demodulation, [], [f9195, f5126])).
fof(f9195, plain, (~ (op2(e21, e24) = h3(e13)) | (~ spl280_79 | spl280_682)), inference(forward_demodulation, [], [f5846, f2244])).
fof(f5846, plain, (~ (op2(e21, e24) = h3(op1(e11, e14))) | spl280_682), inference(avatar_component_clause, [], [f5844])).
fof(f5844, plain, (spl280_682 <=> (op2(e21, e24) = h3(op1(e11, e14)))), introduced(avatar_definition, [new_symbols(naming, [spl280_682])])).
fof(f9156, plain, (~ spl280_109 | ~ spl280_239 | ~ spl280_562 | ~ spl280_572 | spl280_688), inference(avatar_contradiction_clause, [], [f9155])).
fof(f9155, plain, ($false | (~ spl280_109 | ~ spl280_239 | ~ spl280_562 | ~ spl280_572 | spl280_688)), inference(subsumption_resolution, [], [f9154, f2966])).
fof(f2966, plain, ((e23 = op2(e20, e23)) | ~ spl280_239), inference(avatar_component_clause, [], [f2964])).
fof(f2964, plain, (spl280_239 <=> (e23 = op2(e20, e23))), introduced(avatar_definition, [new_symbols(naming, [spl280_239])])).
fof(f9154, plain, (~ (e23 = op2(e20, e23)) | (~ spl280_109 | ~ spl280_562 | ~ spl280_572 | spl280_688)), inference(forward_demodulation, [], [f9153, f5126])).
fof(f9153, plain, (~ (op2(e20, e23) = h3(e13)) | (~ spl280_109 | ~ spl280_562 | ~ spl280_572 | spl280_688)), inference(forward_demodulation, [], [f9152, f2370])).
fof(f2370, plain, ((e13 = op1(e10, e13)) | ~ spl280_109), inference(avatar_component_clause, [], [f2368])).
fof(f2368, plain, (spl280_109 <=> (e13 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl280_109])])).
fof(f9152, plain, (~ (op2(e20, e23) = h3(op1(e10, e13))) | (~ spl280_562 | ~ spl280_572 | spl280_688)), inference(forward_demodulation, [], [f9151, f5190])).
fof(f9151, plain, (~ (h3(op1(e10, e13)) = op2(h2(e13), e23)) | (~ spl280_562 | spl280_688)), inference(forward_demodulation, [], [f5870, f5126])).
fof(f5870, plain, (~ (h3(op1(e10, e13)) = op2(h2(e13), h3(e13))) | spl280_688), inference(avatar_component_clause, [], [f5868])).
fof(f5868, plain, (spl280_688 <=> (h3(op1(e10, e13)) = op2(h2(e13), h3(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl280_688])])).
fof(f9133, plain, (~ spl280_26 | ~ spl280_156 | ~ spl280_562 | ~ spl280_572 | spl280_672), inference(avatar_contradiction_clause, [], [f9132])).
fof(f9132, plain, ($false | (~ spl280_26 | ~ spl280_156 | ~ spl280_562 | ~ spl280_572 | spl280_672)), inference(subsumption_resolution, [], [f9131, f2618])).
fof(f2618, plain, ((e20 = op2(e23, e24)) | ~ spl280_156), inference(avatar_component_clause, [], [f2616])).
fof(f2616, plain, (spl280_156 <=> (e20 = op2(e23, e24))), introduced(avatar_definition, [new_symbols(naming, [spl280_156])])).
fof(f9131, plain, (~ (e20 = op2(e23, e24)) | (~ spl280_26 | ~ spl280_562 | ~ spl280_572 | spl280_672)), inference(forward_demodulation, [], [f9130, f7170])).
fof(f9130, plain, (~ (op2(e23, e24) = h3(e10)) | (~ spl280_26 | ~ spl280_562 | spl280_672)), inference(forward_demodulation, [], [f9129, f2022])).
fof(f9129, plain, (~ (op2(e23, e24) = h3(op1(e13, e14))) | (~ spl280_562 | spl280_672)), inference(forward_demodulation, [], [f5806, f5126])).
fof(f5806, plain, (~ (h3(op1(e13, e14)) = op2(h3(e13), e24)) | spl280_672), inference(avatar_component_clause, [], [f5804])).
fof(f5804, plain, (spl280_672 <=> (h3(op1(e13, e14)) = op2(h3(e13), e24))), introduced(avatar_definition, [new_symbols(naming, [spl280_672])])).
fof(f9102, plain, (~ spl280_586 | ~ spl280_121 | ~ spl280_572 | spl280_691), inference(avatar_split_clause, [], [f9101, f5880, f5189, f2419, f5262])).
fof(f5262, plain, (spl280_586 <=> (e20 = h2(e14))), introduced(avatar_definition, [new_symbols(naming, [spl280_586])])).
fof(f2419, plain, (spl280_121 <=> (e10 = op1(e10, e10))), introduced(avatar_definition, [new_symbols(naming, [spl280_121])])).
fof(f5880, plain, (spl280_691 <=> (h2(e14) = h3(op1(e10, e10)))), introduced(avatar_definition, [new_symbols(naming, [spl280_691])])).
fof(f9101, plain, (~ (e20 = h2(e14)) | (~ spl280_121 | ~ spl280_572 | spl280_691)), inference(forward_demodulation, [], [f9100, f7170])).
fof(f9100, plain, (~ (h2(e14) = h3(e10)) | (~ spl280_121 | spl280_691)), inference(forward_demodulation, [], [f5882, f2421])).
fof(f2421, plain, ((e10 = op1(e10, e10)) | ~ spl280_121), inference(avatar_component_clause, [], [f2419])).
fof(f5882, plain, (~ (h2(e14) = h3(op1(e10, e10))) | spl280_691), inference(avatar_component_clause, [], [f5880])).
fof(f9098, plain, (spl280_528 | ~ spl280_547 | ~ spl280_664), inference(avatar_split_clause, [], [f8828, f5698, f5049, f4953])).
fof(f5049, plain, (spl280_547 <=> (e22 = h4(e14))), introduced(avatar_definition, [new_symbols(naming, [spl280_547])])).
fof(f5698, plain, (spl280_664 <=> (e24 = h4(e13))), introduced(avatar_definition, [new_symbols(naming, [spl280_664])])).
fof(f8828, plain, ((e22 = h5(e13)) | (~ spl280_547 | ~ spl280_664)), inference(backward_demodulation, [], [f7833, f5050])).
fof(f5050, plain, ((e22 = h4(e14)) | ~ spl280_547), inference(avatar_component_clause, [], [f5049])).
fof(f7833, plain, ((h4(e14) = h5(e13)) | ~ spl280_664), inference(backward_demodulation, [], [f1787, f7830])).
fof(f7830, plain, ((op2(e24, e24) = h4(e14)) | ~ spl280_664), inference(backward_demodulation, [], [f4902, f5699])).
fof(f5699, plain, ((e24 = h4(e13)) | ~ spl280_664), inference(avatar_component_clause, [], [f5698])).
fof(f4902, plain, (h4(e14) = op2(h4(e13), h4(e13))), inference(backward_demodulation, [], [f1783, f1782])).
fof(f1782, plain, (op2(e23, e23) = h4(e13)), inference(cnf_transformation, [], [f19])).
fof(f19, plain, ((h4(e14) = op2(op2(e23, e23), op2(e23, e23))) & (op2(e23, e23) = h4(e13)) & (op2(op2(e23, e23), e23) = h4(e11)) & (h4(e10) = op2(op2(op2(e23, e23), e23), op2(op2(e23, e23), e23))) & (e23 = h4(e12))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG051+1.p', ax19)).
fof(f1783, plain, (h4(e14) = op2(op2(e23, e23), op2(e23, e23))), inference(cnf_transformation, [], [f19])).
fof(f1787, plain, (op2(e24, e24) = h5(e13)), inference(cnf_transformation, [], [f20])).
fof(f20, plain, ((h5(e14) = op2(op2(e24, e24), op2(e24, e24))) & (op2(e24, e24) = h5(e13)) & (op2(op2(e24, e24), e24) = h5(e11)) & (h5(e10) = op2(op2(op2(e24, e24), e24), op2(op2(e24, e24), e24))) & (e24 = h5(e12))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG051+1.p', ax20)).
fof(f9067, plain, (~ spl280_159 | ~ spl280_179), inference(avatar_split_clause, [], [f8124, f2712, f2628])).
fof(f2628, plain, (spl280_159 <=> (e23 = op2(e23, e24))), introduced(avatar_definition, [new_symbols(naming, [spl280_159])])).
fof(f8124, plain, (~ (e23 = op2(e23, e24)) | ~ spl280_179), inference(forward_demodulation, [], [f942, f2714])).
fof(f942, plain, ~ (op2(e23, e20) = op2(e23, e24)), inference(cnf_transformation, [], [f8])).
fof(f9032, plain, (~ spl280_6 | ~ spl280_11), inference(avatar_split_clause, [], [f9027, f1957, f1936])).
fof(f1936, plain, (spl280_6 <=> (e10 = op1(e14, e13))), introduced(avatar_definition, [new_symbols(naming, [spl280_6])])).
fof(f9027, plain, (~ (e10 = op1(e14, e13)) | ~ spl280_11), inference(backward_demodulation, [], [f856, f1959])).
fof(f856, plain, ~ (op1(e14, e12) = op1(e14, e13)), inference(cnf_transformation, [], [f7])).
fof(f9018, plain, (~ spl280_29 | ~ spl280_49), inference(avatar_split_clause, [], [f9016, f2116, f2032])).
fof(f2032, plain, (spl280_29 <=> (e13 = op1(e13, e14))), introduced(avatar_definition, [new_symbols(naming, [spl280_29])])).
fof(f9016, plain, (~ (e13 = op1(e13, e14)) | ~ spl280_49), inference(backward_demodulation, [], [f842, f2118])).
fof(f842, plain, ~ (op1(e13, e10) = op1(e13, e14)), inference(cnf_transformation, [], [f7])).
fof(f8982, plain, (~ spl280_77 | ~ spl280_97), inference(avatar_split_clause, [], [f8975, f2318, f2234])).
fof(f2234, plain, (spl280_77 <=> (e11 = op1(e11, e14))), introduced(avatar_definition, [new_symbols(naming, [spl280_77])])).
fof(f8975, plain, (~ (e11 = op1(e11, e14)) | ~ spl280_97), inference(backward_demodulation, [], [f822, f2320])).
fof(f822, plain, ~ (op1(e11, e10) = op1(e11, e14)), inference(cnf_transformation, [], [f7])).
fof(f8981, plain, (~ spl280_82 | ~ spl280_97), inference(avatar_split_clause, [], [f8974, f2318, f2255])).
fof(f2255, plain, (spl280_82 <=> (e11 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl280_82])])).
fof(f8974, plain, (~ (e11 = op1(e11, e13)) | ~ spl280_97), inference(backward_demodulation, [], [f821, f2320])).
fof(f821, plain, ~ (op1(e11, e10) = op1(e11, e13)), inference(cnf_transformation, [], [f7])).
fof(f8961, plain, (~ spl280_111 | ~ spl280_121), inference(avatar_split_clause, [], [f8953, f2419, f2377])).
fof(f2377, plain, (spl280_111 <=> (e10 = op1(e10, e12))), introduced(avatar_definition, [new_symbols(naming, [spl280_111])])).
fof(f8953, plain, (~ (e10 = op1(e10, e12)) | ~ spl280_121), inference(backward_demodulation, [], [f810, f2421])).
fof(f810, plain, ~ (op1(e10, e10) = op1(e10, e12)), inference(cnf_transformation, [], [f7])).
fof(f8960, plain, (~ spl280_46 | ~ spl280_121), inference(avatar_split_clause, [], [f8951, f2419, f2104])).
fof(f2104, plain, (spl280_46 <=> (e10 = op1(e13, e10))), introduced(avatar_definition, [new_symbols(naming, [spl280_46])])).
fof(f8951, plain, (~ (e10 = op1(e13, e10)) | ~ spl280_121), inference(backward_demodulation, [], [f761, f2421])).
fof(f761, plain, ~ (op1(e10, e10) = op1(e13, e10)), inference(cnf_transformation, [], [f7])).
fof(f8948, plain, (spl280_49 | ~ spl280_126), inference(avatar_split_clause, [], [f8935, f2440, f2116])).
fof(f2440, plain, (spl280_126 <=> (e10 = unit1)), introduced(avatar_definition, [new_symbols(naming, [spl280_126])])).
fof(f8935, plain, ((e13 = op1(e13, e10)) | ~ spl280_126), inference(backward_demodulation, [], [f619, f2442])).
fof(f2442, plain, ((e10 = unit1) | ~ spl280_126), inference(avatar_component_clause, [], [f2440])).
fof(f619, plain, (e13 = op1(e13, unit1)), inference(cnf_transformation, [], [f2])).
fof(f2, plain, (((e14 = unit1) | (e13 = unit1) | (e12 = unit1) | (e11 = unit1) | (e10 = unit1)) & (e14 = op1(e14, unit1)) & (e14 = op1(unit1, e14)) & (e13 = op1(e13, unit1)) & (e13 = op1(unit1, e13)) & (e12 = op1(e12, unit1)) & (e12 = op1(unit1, e12)) & (e11 = op1(e11, unit1)) & (e11 = op1(unit1, e11)) & (e10 = op1(e10, unit1)) & (e10 = op1(unit1, e10))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG051+1.p', ax2)).
fof(f8947, plain, (spl280_73 | ~ spl280_126), inference(avatar_split_clause, [], [f8933, f2440, f2217])).
fof(f8933, plain, ((e12 = op1(e12, e10)) | ~ spl280_126), inference(backward_demodulation, [], [f617, f2442])).
fof(f617, plain, (e12 = op1(e12, unit1)), inference(cnf_transformation, [], [f2])).
fof(f8946, plain, (spl280_113 | ~ spl280_126), inference(avatar_split_clause, [], [f8932, f2440, f2385])).
fof(f8932, plain, ((e12 = op1(e10, e12)) | ~ spl280_126), inference(backward_demodulation, [], [f616, f2442])).
fof(f616, plain, (e12 = op1(unit1, e12)), inference(cnf_transformation, [], [f2])).
fof(f8945, plain, (spl280_97 | ~ spl280_126), inference(avatar_split_clause, [], [f8931, f2440, f2318])).
fof(f8931, plain, ((e11 = op1(e11, e10)) | ~ spl280_126), inference(backward_demodulation, [], [f615, f2442])).
fof(f615, plain, (e11 = op1(e11, unit1)), inference(cnf_transformation, [], [f2])).
fof(f8944, plain, (spl280_117 | ~ spl280_126), inference(avatar_split_clause, [], [f8930, f2440, f2402])).
fof(f8930, plain, ((e11 = op1(e10, e11)) | ~ spl280_126), inference(backward_demodulation, [], [f614, f2442])).
fof(f614, plain, (e11 = op1(unit1, e11)), inference(cnf_transformation, [], [f2])).
fof(f8943, plain, (spl280_121 | ~ spl280_126), inference(avatar_split_clause, [], [f8929, f2440, f2419])).
fof(f8929, plain, ((e10 = op1(e10, e10)) | ~ spl280_126), inference(backward_demodulation, [], [f613, f2442])).
fof(f613, plain, (e10 = op1(e10, unit1)), inference(cnf_transformation, [], [f2])).
fof(f8917, plain, (~ spl280_166 | ~ spl280_167), inference(avatar_contradiction_clause, [], [f8916])).
fof(f8916, plain, ($false | (~ spl280_166 | ~ spl280_167)), inference(subsumption_resolution, [], [f8915, f969])).
fof(f969, plain, ~ (e20 = e21), inference(cnf_transformation, [], [f10])).
fof(f8915, plain, ((e20 = e21) | (~ spl280_166 | ~ spl280_167)), inference(backward_demodulation, [], [f2664, f2660])).
fof(f2660, plain, ((e20 = op2(e23, e22)) | ~ spl280_166), inference(avatar_component_clause, [], [f2658])).
fof(f2658, plain, (spl280_166 <=> (e20 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl280_166])])).
fof(f2664, plain, ((e21 = op2(e23, e22)) | ~ spl280_167), inference(avatar_component_clause, [], [f2662])).
fof(f2662, plain, (spl280_167 <=> (e21 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl280_167])])).
fof(f8884, plain, (~ spl280_214 | ~ spl280_239), inference(avatar_split_clause, [], [f8882, f2964, f2859])).
fof(f2859, plain, (spl280_214 <=> (e23 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl280_214])])).
fof(f8882, plain, (~ (e23 = op2(e21, e23)) | ~ spl280_239), inference(backward_demodulation, [], [f889, f2966])).
fof(f889, plain, ~ (op2(e20, e23) = op2(e21, e23)), inference(cnf_transformation, [], [f8])).
fof(f8863, plain, (spl280_239 | ~ spl280_256), inference(avatar_split_clause, [], [f8849, f3036, f2964])).
fof(f3036, plain, (spl280_256 <=> (e20 = unit2)), introduced(avatar_definition, [new_symbols(naming, [spl280_256])])).
fof(f8849, plain, ((e23 = op2(e20, e23)) | ~ spl280_256), inference(backward_demodulation, [], [f704, f3038])).
fof(f3038, plain, ((e20 = unit2) | ~ spl280_256), inference(avatar_component_clause, [], [f3036])).
fof(f704, plain, (e23 = op2(unit2, e23)), inference(cnf_transformation, [], [f5])).
fof(f5, plain, (((e24 = unit2) | (e23 = unit2) | (e22 = unit2) | (e21 = unit2) | (e20 = unit2)) & (e24 = op2(e24, unit2)) & (e24 = op2(unit2, e24)) & (e23 = op2(e23, unit2)) & (e23 = op2(unit2, e23)) & (e22 = op2(e22, unit2)) & (e22 = op2(unit2, e22)) & (e21 = op2(e21, unit2)) & (e21 = op2(unit2, e21)) & (e20 = op2(e20, unit2)) & (e20 = op2(unit2, e20))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG051+1.p', ax5)).
fof(f8862, plain, (spl280_203 | ~ spl280_256), inference(avatar_split_clause, [], [f8848, f3036, f2813])).
fof(f8848, plain, ((e22 = op2(e22, e20)) | ~ spl280_256), inference(backward_demodulation, [], [f703, f3038])).
fof(f703, plain, (e22 = op2(e22, unit2)), inference(cnf_transformation, [], [f5])).
fof(f8861, plain, (spl280_243 | ~ spl280_256), inference(avatar_split_clause, [], [f8847, f3036, f2981])).
fof(f8847, plain, ((e22 = op2(e20, e22)) | ~ spl280_256), inference(backward_demodulation, [], [f702, f3038])).
fof(f702, plain, (e22 = op2(unit2, e22)), inference(cnf_transformation, [], [f5])).
fof(f8860, plain, (spl280_227 | ~ spl280_256), inference(avatar_split_clause, [], [f8846, f3036, f2914])).
fof(f8846, plain, ((e21 = op2(e21, e20)) | ~ spl280_256), inference(backward_demodulation, [], [f701, f3038])).
fof(f701, plain, (e21 = op2(e21, unit2)), inference(cnf_transformation, [], [f5])).
fof(f8859, plain, (spl280_247 | ~ spl280_256), inference(avatar_split_clause, [], [f8845, f3036, f2998])).
fof(f2998, plain, (spl280_247 <=> (e21 = op2(e20, e21))), introduced(avatar_definition, [new_symbols(naming, [spl280_247])])).
fof(f8845, plain, ((e21 = op2(e20, e21)) | ~ spl280_256), inference(backward_demodulation, [], [f700, f3038])).
fof(f700, plain, (e21 = op2(unit2, e21)), inference(cnf_transformation, [], [f5])).
fof(f8858, plain, (spl280_251 | ~ spl280_256), inference(avatar_split_clause, [], [f8844, f3036, f3015])).
fof(f3015, plain, (spl280_251 <=> (e20 = op2(e20, e20))), introduced(avatar_definition, [new_symbols(naming, [spl280_251])])).
fof(f8844, plain, ((e20 = op2(e20, e20)) | ~ spl280_256), inference(backward_demodulation, [], [f699, f3038])).
fof(f699, plain, (e20 = op2(e20, unit2)), inference(cnf_transformation, [], [f5])).
fof(f8800, plain, (spl280_586 | ~ spl280_572 | ~ spl280_606), inference(avatar_split_clause, [], [f8795, f5363, f5189, f5262])).
fof(f5363, plain, (spl280_606 <=> (e20 = h1(e13))), introduced(avatar_definition, [new_symbols(naming, [spl280_606])])).
fof(f8795, plain, ((e20 = h2(e14)) | (~ spl280_572 | ~ spl280_606)), inference(backward_demodulation, [], [f7178, f5364])).
fof(f5364, plain, ((e20 = h1(e13)) | ~ spl280_606), inference(avatar_component_clause, [], [f5363])).
fof(f7178, plain, ((h1(e13) = h2(e14)) | ~ spl280_572), inference(forward_demodulation, [], [f7168, f1767])).
fof(f1767, plain, (op2(e20, e20) = h1(e13)), inference(cnf_transformation, [], [f16])).
fof(f16, plain, ((h1(e14) = op2(op2(e20, e20), op2(e20, e20))) & (op2(e20, e20) = h1(e13)) & (op2(op2(e20, e20), e20) = h1(e11)) & (h1(e10) = op2(op2(op2(e20, e20), e20), op2(op2(e20, e20), e20))) & (e20 = h1(e12))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG051+1.p', ax16)).
fof(f7168, plain, ((op2(e20, e20) = h2(e14)) | ~ spl280_572), inference(backward_demodulation, [], [f4874, f5190])).
fof(f4874, plain, (h2(e14) = op2(h2(e13), h2(e13))), inference(backward_demodulation, [], [f1773, f1772])).
fof(f1773, plain, (h2(e14) = op2(op2(e21, e21), op2(e21, e21))), inference(cnf_transformation, [], [f17])).
fof(f8781, plain, (~ spl280_43 | ~ spl280_173 | ~ spl280_562 | spl280_675), inference(avatar_contradiction_clause, [], [f8780])).
fof(f8780, plain, ($false | (~ spl280_43 | ~ spl280_173 | ~ spl280_562 | spl280_675)), inference(subsumption_resolution, [], [f8779, f2689])).
fof(f2689, plain, ((e22 = op2(e23, e21)) | ~ spl280_173), inference(avatar_component_clause, [], [f2687])).
fof(f2687, plain, (spl280_173 <=> (e22 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl280_173])])).
fof(f8779, plain, (~ (e22 = op2(e23, e21)) | (~ spl280_43 | ~ spl280_562 | spl280_675)), inference(forward_demodulation, [], [f8778, f1774])).
fof(f8778, plain, (~ (op2(e23, e21) = h3(e12)) | (~ spl280_43 | ~ spl280_562 | spl280_675)), inference(forward_demodulation, [], [f8777, f2093])).
fof(f2093, plain, ((e12 = op1(e13, e11)) | ~ spl280_43), inference(avatar_component_clause, [], [f2091])).
fof(f2091, plain, (spl280_43 <=> (e12 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl280_43])])).
fof(f8777, plain, (~ (op2(e23, e21) = h3(op1(e13, e11))) | (~ spl280_562 | spl280_675)), inference(forward_demodulation, [], [f5818, f5126])).
fof(f5818, plain, (~ (h3(op1(e13, e11)) = op2(h3(e13), e21)) | spl280_675), inference(avatar_component_clause, [], [f5816])).
fof(f5816, plain, (spl280_675 <=> (h3(op1(e13, e11)) = op2(h3(e13), e21))), introduced(avatar_definition, [new_symbols(naming, [spl280_675])])).
fof(f8720, plain, (~ spl280_117 | ~ spl280_17), inference(avatar_split_clause, [], [f8719, f1982, f2402])).
fof(f1982, plain, (spl280_17 <=> (e11 = op1(e14, e11))), introduced(avatar_definition, [new_symbols(naming, [spl280_17])])).
fof(f8719, plain, (~ (e11 = op1(e10, e11)) | ~ spl280_17), inference(forward_demodulation, [], [f772, f1984])).
fof(f1984, plain, ((e11 = op1(e14, e11)) | ~ spl280_17), inference(avatar_component_clause, [], [f1982])).
fof(f772, plain, ~ (op1(e10, e11) = op1(e14, e11)), inference(cnf_transformation, [], [f7])).
fof(f8692, plain, (~ spl280_24 | ~ spl280_25), inference(avatar_contradiction_clause, [], [f8691])).
fof(f8691, plain, ($false | (~ spl280_24 | ~ spl280_25)), inference(subsumption_resolution, [], [f8690, f968])).
fof(f968, plain, ~ (e13 = e14), inference(cnf_transformation, [], [f9])).
fof(f8690, plain, ((e13 = e14) | (~ spl280_24 | ~ spl280_25)), inference(backward_demodulation, [], [f2017, f2013])).
fof(f2013, plain, ((e13 = op1(e14, e10)) | ~ spl280_24), inference(avatar_component_clause, [], [f2011])).
fof(f2011, plain, (spl280_24 <=> (e13 = op1(e14, e10))), introduced(avatar_definition, [new_symbols(naming, [spl280_24])])).
fof(f8646, plain, (~ spl280_91 | ~ spl280_94), inference(avatar_contradiction_clause, [], [f8645])).
fof(f8645, plain, ($false | (~ spl280_91 | ~ spl280_94)), inference(subsumption_resolution, [], [f8644, f961])).
fof(f961, plain, ~ (e10 = e13), inference(cnf_transformation, [], [f9])).
fof(f8644, plain, ((e10 = e13) | (~ spl280_91 | ~ spl280_94)), inference(forward_demodulation, [], [f2307, f2295])).
fof(f2307, plain, ((e13 = op1(e11, e11)) | ~ spl280_94), inference(avatar_component_clause, [], [f2305])).
fof(f2305, plain, (spl280_94 <=> (e13 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl280_94])])).
fof(f8596, plain, (spl280_13 | ~ spl280_130), inference(avatar_split_clause, [], [f8585, f2456, f1965])).
fof(f1965, plain, (spl280_13 <=> (e12 = op1(e14, e12))), introduced(avatar_definition, [new_symbols(naming, [spl280_13])])).
fof(f2456, plain, (spl280_130 <=> (e14 = unit1)), introduced(avatar_definition, [new_symbols(naming, [spl280_130])])).
fof(f8585, plain, ((e12 = op1(e14, e12)) | ~ spl280_130), inference(backward_demodulation, [], [f616, f2458])).
fof(f2458, plain, ((e14 = unit1) | ~ spl280_130), inference(avatar_component_clause, [], [f2456])).
fof(f8594, plain, (spl280_101 | ~ spl280_130), inference(avatar_contradiction_clause, [], [f8593])).
fof(f8593, plain, ($false | (spl280_101 | ~ spl280_130)), inference(subsumption_resolution, [], [f8582, f2336])).
fof(f2336, plain, (~ (e10 = op1(e10, e14)) | spl280_101), inference(avatar_component_clause, [], [f2335])).
fof(f2335, plain, (spl280_101 <=> (e10 = op1(e10, e14))), introduced(avatar_definition, [new_symbols(naming, [spl280_101])])).
fof(f8582, plain, ((e10 = op1(e10, e14)) | ~ spl280_130), inference(backward_demodulation, [], [f613, f2458])).
fof(f8575, plain, (~ spl280_141 | ~ spl280_143), inference(avatar_contradiction_clause, [], [f8574])).
fof(f8574, plain, ($false | (~ spl280_141 | ~ spl280_143)), inference(subsumption_resolution, [], [f8572, f970])).
fof(f970, plain, ~ (e20 = e22), inference(cnf_transformation, [], [f10])).
fof(f8572, plain, ((e20 = e22) | (~ spl280_141 | ~ spl280_143)), inference(backward_demodulation, [], [f2563, f2555])).
fof(f2563, plain, ((e22 = op2(e24, e22)) | ~ spl280_143), inference(avatar_component_clause, [], [f2561])).
fof(f2561, plain, (spl280_143 <=> (e22 = op2(e24, e22))), introduced(avatar_definition, [new_symbols(naming, [spl280_143])])).
fof(f8532, plain, (~ spl280_221 | ~ spl280_225), inference(avatar_contradiction_clause, [], [f8531])).
fof(f8531, plain, ($false | (~ spl280_221 | ~ spl280_225)), inference(subsumption_resolution, [], [f8530, f972])).
fof(f972, plain, ~ (e20 = e24), inference(cnf_transformation, [], [f10])).
fof(f8530, plain, ((e20 = e24) | (~ spl280_221 | ~ spl280_225)), inference(forward_demodulation, [], [f2907, f2891])).
fof(f2907, plain, ((e24 = op2(e21, e21)) | ~ spl280_225), inference(avatar_component_clause, [], [f2905])).
fof(f2905, plain, (spl280_225 <=> (e24 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl280_225])])).
fof(f8461, plain, (spl280_231 | ~ spl280_260), inference(avatar_split_clause, [], [f8191, f3052, f2931])).
fof(f2931, plain, (spl280_231 <=> (e20 = op2(e20, e24))), introduced(avatar_definition, [new_symbols(naming, [spl280_231])])).
fof(f3052, plain, (spl280_260 <=> (e24 = unit2)), introduced(avatar_definition, [new_symbols(naming, [spl280_260])])).
fof(f8191, plain, ((e20 = op2(e20, e24)) | ~ spl280_260), inference(backward_demodulation, [], [f699, f3054])).
fof(f3054, plain, ((e24 = unit2) | ~ spl280_260), inference(avatar_component_clause, [], [f3052])).
fof(f8436, plain, (~ spl280_199 | ~ spl280_562), inference(avatar_split_clause, [], [f7186, f5125, f2796])).
fof(f2796, plain, (spl280_199 <=> (e23 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl280_199])])).
fof(f7186, plain, (~ (e23 = op2(e22, e21)) | ~ spl280_562), inference(backward_demodulation, [], [f4884, f5126])).
fof(f4884, plain, ~ (op2(e22, e21) = h3(e13)), inference(backward_demodulation, [], [f933, f1777])).
fof(f933, plain, ~ (op2(e22, e21) = op2(e22, e22)), inference(cnf_transformation, [], [f8])).
fof(f8421, plain, (~ spl280_175 | ~ spl280_664), inference(avatar_split_clause, [], [f8120, f5698, f2695])).
fof(f2695, plain, (spl280_175 <=> (e24 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl280_175])])).
fof(f8120, plain, (~ (e24 = op2(e23, e21)) | ~ spl280_664), inference(forward_demodulation, [], [f4899, f5699])).
fof(f4899, plain, ~ (op2(e23, e21) = h4(e13)), inference(backward_demodulation, [], [f944, f1782])).
fof(f944, plain, ~ (op2(e23, e21) = op2(e23, e23)), inference(cnf_transformation, [], [f8])).
fof(f8397, plain, (spl280_143 | ~ spl280_260), inference(avatar_split_clause, [], [f8194, f3052, f2561])).
fof(f8194, plain, ((e22 = op2(e24, e22)) | ~ spl280_260), inference(backward_demodulation, [], [f702, f3054])).
fof(f8388, plain, (~ spl280_120 | ~ spl280_105), inference(avatar_split_clause, [], [f8387, f2351, f2414])).
fof(f2414, plain, (spl280_120 <=> (e14 = op1(e10, e11))), introduced(avatar_definition, [new_symbols(naming, [spl280_120])])).
fof(f8387, plain, (~ (e14 = op1(e10, e11)) | ~ spl280_105), inference(forward_demodulation, [], [f815, f2353])).
fof(f815, plain, ~ (op1(e10, e11) = op1(e10, e14)), inference(cnf_transformation, [], [f7])).
fof(f8386, plain, (~ spl280_100 | ~ spl280_25), inference(avatar_split_clause, [], [f8385, f2015, f2330])).
fof(f2330, plain, (spl280_100 <=> (e14 = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl280_100])])).
fof(f8385, plain, (~ (e14 = op1(e11, e10)) | ~ spl280_25), inference(forward_demodulation, [], [f765, f2017])).
fof(f765, plain, ~ (op1(e11, e10) = op1(e14, e10)), inference(cnf_transformation, [], [f7])).
fof(f8370, plain, (~ spl280_20 | ~ spl280_25), inference(avatar_split_clause, [], [f8369, f2015, f1994])).
fof(f1994, plain, (spl280_20 <=> (e14 = op1(e14, e11))), introduced(avatar_definition, [new_symbols(naming, [spl280_20])])).
fof(f8369, plain, (~ (e14 = op1(e14, e11)) | ~ spl280_25), inference(forward_demodulation, [], [f849, f2017])).
fof(f849, plain, ~ (op1(e14, e10) = op1(e14, e11)), inference(cnf_transformation, [], [f7])).
fof(f8356, plain, (~ spl280_31 | ~ spl280_35), inference(avatar_contradiction_clause, [], [f8355])).
fof(f8355, plain, ($false | (~ spl280_31 | ~ spl280_35)), inference(subsumption_resolution, [], [f8353, f962])).
fof(f962, plain, ~ (e10 = e14), inference(cnf_transformation, [], [f9])).
fof(f8353, plain, ((e10 = e14) | (~ spl280_31 | ~ spl280_35)), inference(backward_demodulation, [], [f2059, f2043])).
fof(f2043, plain, ((e10 = op1(e13, e13)) | ~ spl280_31), inference(avatar_component_clause, [], [f2041])).
fof(f2041, plain, (spl280_31 <=> (e10 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl280_31])])).
fof(f8343, plain, (~ spl280_61 | ~ spl280_64), inference(avatar_contradiction_clause, [], [f8342])).
fof(f8342, plain, ($false | (~ spl280_61 | ~ spl280_64)), inference(subsumption_resolution, [], [f8341, f961])).
fof(f8341, plain, ((e10 = e13) | (~ spl280_61 | ~ spl280_64)), inference(backward_demodulation, [], [f2181, f2169])).
fof(f2169, plain, ((e10 = op1(e12, e12)) | ~ spl280_61), inference(avatar_component_clause, [], [f2167])).
fof(f2167, plain, (spl280_61 <=> (e10 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl280_61])])).
fof(f8323, plain, (~ spl280_91 | ~ spl280_95), inference(avatar_contradiction_clause, [], [f8322])).
fof(f8322, plain, ($false | (~ spl280_91 | ~ spl280_95)), inference(subsumption_resolution, [], [f8321, f962])).
fof(f8321, plain, ((e10 = e14) | (~ spl280_91 | ~ spl280_95)), inference(forward_demodulation, [], [f2311, f2295])).
fof(f2311, plain, ((e14 = op1(e11, e11)) | ~ spl280_95), inference(avatar_component_clause, [], [f2309])).
fof(f2309, plain, (spl280_95 <=> (e14 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl280_95])])).
fof(f8307, plain, (spl280_25 | ~ spl280_126), inference(avatar_split_clause, [], [f8298, f2440, f2015])).
fof(f8298, plain, ((e14 = op1(e14, e10)) | ~ spl280_126), inference(backward_demodulation, [], [f621, f2442])).
fof(f621, plain, (e14 = op1(e14, unit1)), inference(cnf_transformation, [], [f2])).
fof(f8254, plain, (~ spl280_192 | ~ spl280_194), inference(avatar_contradiction_clause, [], [f8253])).
fof(f8253, plain, ($false | (~ spl280_192 | ~ spl280_194)), inference(subsumption_resolution, [], [f8252, f974])).
fof(f974, plain, ~ (e21 = e23), inference(cnf_transformation, [], [f10])).
fof(f8252, plain, ((e21 = e23) | (~ spl280_192 | ~ spl280_194)), inference(backward_demodulation, [], [f2777, f2769])).
fof(f2769, plain, ((e21 = op2(e22, e22)) | ~ spl280_192), inference(avatar_component_clause, [], [f2767])).
fof(f2767, plain, (spl280_192 <=> (e21 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl280_192])])).
fof(f2777, plain, ((e23 = op2(e22, e22)) | ~ spl280_194), inference(avatar_component_clause, [], [f2775])).
fof(f2775, plain, (spl280_194 <=> (e23 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl280_194])])).
fof(f8151, plain, (~ spl280_227 | ~ spl280_212), inference(avatar_split_clause, [], [f8150, f2851, f2914])).
fof(f2851, plain, (spl280_212 <=> (e21 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl280_212])])).
fof(f8150, plain, (~ (e21 = op2(e21, e20)) | ~ spl280_212), inference(forward_demodulation, [], [f921, f2853])).
fof(f2853, plain, ((e21 = op2(e21, e23)) | ~ spl280_212), inference(avatar_component_clause, [], [f2851])).
fof(f921, plain, ~ (op2(e21, e20) = op2(e21, e23)), inference(cnf_transformation, [], [f8])).
fof(f8137, plain, (~ spl280_203 | ~ spl280_188), inference(avatar_split_clause, [], [f8136, f2750, f2813])).
fof(f2750, plain, (spl280_188 <=> (e22 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl280_188])])).
fof(f8136, plain, (~ (e22 = op2(e22, e20)) | ~ spl280_188), inference(forward_demodulation, [], [f931, f2752])).
fof(f2752, plain, ((e22 = op2(e22, e23)) | ~ spl280_188), inference(avatar_component_clause, [], [f2750])).
fof(f931, plain, ~ (op2(e22, e20) = op2(e22, e23)), inference(cnf_transformation, [], [f8])).
fof(f8089, plain, (spl280_547 | ~ spl280_133 | ~ spl280_664), inference(avatar_split_clause, [], [f8088, f5698, f2519, f5049])).
fof(f2519, plain, (spl280_133 <=> (e22 = op2(e24, e24))), introduced(avatar_definition, [new_symbols(naming, [spl280_133])])).
fof(f8088, plain, ((e22 = h4(e14)) | (~ spl280_133 | ~ spl280_664)), inference(forward_demodulation, [], [f7830, f2521])).
fof(f2521, plain, ((e22 = op2(e24, e24)) | ~ spl280_133), inference(avatar_component_clause, [], [f2519])).
fof(f8046, plain, (~ spl280_49 | ~ spl280_99), inference(avatar_split_clause, [], [f8045, f2326, f2116])).
fof(f2326, plain, (spl280_99 <=> (e13 = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl280_99])])).
fof(f8045, plain, (~ (e13 = op1(e13, e10)) | ~ spl280_99), inference(forward_demodulation, [], [f764, f2328])).
fof(f2328, plain, ((e13 = op1(e11, e10)) | ~ spl280_99), inference(avatar_component_clause, [], [f2326])).
fof(f764, plain, ~ (op1(e11, e10) = op1(e13, e10)), inference(cnf_transformation, [], [f7])).
fof(f8013, plain, (~ spl280_36 | ~ spl280_37), inference(avatar_contradiction_clause, [], [f8012])).
fof(f8012, plain, ($false | (~ spl280_36 | ~ spl280_37)), inference(subsumption_resolution, [], [f8011, f959])).
fof(f8011, plain, ((e10 = e11) | (~ spl280_36 | ~ spl280_37)), inference(backward_demodulation, [], [f2068, f2064])).
fof(f2064, plain, ((e10 = op1(e13, e12)) | ~ spl280_36), inference(avatar_component_clause, [], [f2062])).
fof(f2062, plain, (spl280_36 <=> (e10 = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl280_36])])).
fof(f7965, plain, (~ spl280_37 | ~ spl280_129), inference(avatar_contradiction_clause, [], [f7964])).
fof(f7964, plain, ($false | (~ spl280_37 | ~ spl280_129)), inference(subsumption_resolution, [], [f7963, f963])).
fof(f963, plain, ~ (e11 = e12), inference(cnf_transformation, [], [f9])).
fof(f7963, plain, ((e11 = e12) | (~ spl280_37 | ~ spl280_129)), inference(forward_demodulation, [], [f7951, f2068])).
fof(f7951, plain, ((e12 = op1(e13, e12)) | ~ spl280_129), inference(backward_demodulation, [], [f616, f2454])).
fof(f2454, plain, ((e13 = unit1) | ~ spl280_129), inference(avatar_component_clause, [], [f2452])).
fof(f2452, plain, (spl280_129 <=> (e13 = unit1)), introduced(avatar_definition, [new_symbols(naming, [spl280_129])])).
fof(f7922, plain, (~ spl280_176 | ~ spl280_179), inference(avatar_contradiction_clause, [], [f7921])).
fof(f7921, plain, ($false | (~ spl280_176 | ~ spl280_179)), inference(subsumption_resolution, [], [f7920, f971])).
fof(f7920, plain, ((e20 = e23) | (~ spl280_176 | ~ spl280_179)), inference(forward_demodulation, [], [f2714, f2702])).
fof(f2702, plain, ((e20 = op2(e23, e20)) | ~ spl280_176), inference(avatar_component_clause, [], [f2700])).
fof(f2700, plain, (spl280_176 <=> (e20 = op2(e23, e20))), introduced(avatar_definition, [new_symbols(naming, [spl280_176])])).
fof(f7918, plain, (~ spl280_191 | ~ spl280_194), inference(avatar_contradiction_clause, [], [f7917])).
fof(f7917, plain, ($false | (~ spl280_191 | ~ spl280_194)), inference(subsumption_resolution, [], [f7916, f971])).
fof(f7916, plain, ((e20 = e23) | (~ spl280_191 | ~ spl280_194)), inference(backward_demodulation, [], [f2777, f2765])).
fof(f2765, plain, ((e20 = op2(e22, e22)) | ~ spl280_191), inference(avatar_component_clause, [], [f2763])).
fof(f2763, plain, (spl280_191 <=> (e20 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl280_191])])).
fof(f7897, plain, (~ spl280_227 | ~ spl280_229), inference(avatar_contradiction_clause, [], [f7896])).
fof(f7896, plain, ($false | (~ spl280_227 | ~ spl280_229)), inference(subsumption_resolution, [], [f7895, f974])).
fof(f7895, plain, ((e21 = e23) | (~ spl280_227 | ~ spl280_229)), inference(forward_demodulation, [], [f2924, f2916])).
fof(f2924, plain, ((e23 = op2(e21, e20)) | ~ spl280_229), inference(avatar_component_clause, [], [f2922])).
fof(f2922, plain, (spl280_229 <=> (e23 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl280_229])])).
fof(f7894, plain, (~ spl280_227 | ~ spl280_230), inference(avatar_contradiction_clause, [], [f7893])).
fof(f7893, plain, ($false | (~ spl280_227 | ~ spl280_230)), inference(subsumption_resolution, [], [f7892, f975])).
fof(f975, plain, ~ (e21 = e24), inference(cnf_transformation, [], [f10])).
fof(f7892, plain, ((e21 = e24) | (~ spl280_227 | ~ spl280_230)), inference(forward_demodulation, [], [f2928, f2916])).
fof(f2928, plain, ((e24 = op2(e21, e20)) | ~ spl280_230), inference(avatar_component_clause, [], [f2926])).
fof(f2926, plain, (spl280_230 <=> (e24 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl280_230])])).
fof(f7879, plain, (~ spl280_241 | ~ spl280_243), inference(avatar_contradiction_clause, [], [f7878])).
fof(f7878, plain, ($false | (~ spl280_241 | ~ spl280_243)), inference(subsumption_resolution, [], [f7877, f970])).
fof(f7877, plain, ((e20 = e22) | (~ spl280_241 | ~ spl280_243)), inference(forward_demodulation, [], [f2983, f2975])).
fof(f2975, plain, ((e20 = op2(e20, e22)) | ~ spl280_241), inference(avatar_component_clause, [], [f2973])).
fof(f2973, plain, (spl280_241 <=> (e20 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl280_241])])).
fof(f7861, plain, (~ spl280_167 | ~ spl280_259), inference(avatar_contradiction_clause, [], [f7860])).
fof(f7860, plain, ($false | (~ spl280_167 | ~ spl280_259)), inference(subsumption_resolution, [], [f7859, f973])).
fof(f973, plain, ~ (e21 = e22), inference(cnf_transformation, [], [f10])).
fof(f7859, plain, ((e21 = e22) | (~ spl280_167 | ~ spl280_259)), inference(forward_demodulation, [], [f7847, f2664])).
fof(f7847, plain, ((e22 = op2(e23, e22)) | ~ spl280_259), inference(backward_demodulation, [], [f702, f3050])).
fof(f3050, plain, ((e23 = unit2) | ~ spl280_259), inference(avatar_component_clause, [], [f3048])).
fof(f3048, plain, (spl280_259 <=> (e23 = unit2)), introduced(avatar_definition, [new_symbols(naming, [spl280_259])])).
fof(f7832, plain, (~ spl280_215 | ~ spl280_664), inference(avatar_split_clause, [], [f7828, f5698, f2863])).
fof(f2863, plain, (spl280_215 <=> (e24 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl280_215])])).
fof(f7828, plain, (~ (e24 = op2(e21, e23)) | ~ spl280_664), inference(backward_demodulation, [], [f4895, f5699])).
fof(f4895, plain, ~ (op2(e21, e23) = h4(e13)), inference(backward_demodulation, [], [f894, f1782])).
fof(f894, plain, ~ (op2(e21, e23) = op2(e23, e23)), inference(cnf_transformation, [], [f8])).
fof(f7826, plain, (spl280_664 | ~ spl280_165), inference(avatar_split_clause, [], [f7659, f2653, f5698])).
fof(f2653, plain, (spl280_165 <=> (e24 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl280_165])])).
fof(f7659, plain, ((e24 = h4(e13)) | ~ spl280_165), inference(backward_demodulation, [], [f1782, f2655])).
fof(f2655, plain, ((e24 = op2(e23, e23)) | ~ spl280_165), inference(avatar_component_clause, [], [f2653])).
fof(f7809, plain, (~ spl280_553 | ~ spl280_167), inference(avatar_split_clause, [], [f7808, f2662, f5079])).
fof(f5079, plain, (spl280_553 <=> (e21 = h4(e13))), introduced(avatar_definition, [new_symbols(naming, [spl280_553])])).
fof(f7808, plain, (~ (e21 = h4(e13)) | ~ spl280_167), inference(forward_demodulation, [], [f4900, f2664])).
fof(f4900, plain, ~ (op2(e23, e22) = h4(e13)), inference(backward_demodulation, [], [f946, f1782])).
fof(f946, plain, ~ (op2(e23, e22) = op2(e23, e23)), inference(cnf_transformation, [], [f8])).
fof(f7798, plain, (~ spl280_158 | ~ spl280_173), inference(avatar_split_clause, [], [f7797, f2687, f2624])).
fof(f2624, plain, (spl280_158 <=> (e22 = op2(e23, e24))), introduced(avatar_definition, [new_symbols(naming, [spl280_158])])).
fof(f7797, plain, (~ (e22 = op2(e23, e24)) | ~ spl280_173), inference(forward_demodulation, [], [f945, f2689])).
fof(f945, plain, ~ (op2(e23, e21) = op2(e23, e24)), inference(cnf_transformation, [], [f8])).
fof(f7774, plain, (~ spl280_45 | ~ spl280_35), inference(avatar_split_clause, [], [f7773, f2057, f2099])).
fof(f2099, plain, (spl280_45 <=> (e14 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl280_45])])).
fof(f7773, plain, (~ (e14 = op1(e13, e11)) | ~ spl280_35), inference(forward_demodulation, [], [f844, f2059])).
fof(f844, plain, ~ (op1(e13, e11) = op1(e13, e13)), inference(cnf_transformation, [], [f7])).
fof(f7708, plain, (~ spl280_37 | ~ spl280_128), inference(avatar_contradiction_clause, [], [f7707])).
fof(f7707, plain, ($false | (~ spl280_37 | ~ spl280_128)), inference(subsumption_resolution, [], [f7706, f964])).
fof(f964, plain, ~ (e11 = e13), inference(cnf_transformation, [], [f9])).
fof(f7706, plain, ((e11 = e13) | (~ spl280_37 | ~ spl280_128)), inference(forward_demodulation, [], [f7691, f2068])).
fof(f7691, plain, ((e13 = op1(e13, e12)) | ~ spl280_128), inference(backward_demodulation, [], [f619, f2450])).
fof(f2450, plain, ((e12 = unit1) | ~ spl280_128), inference(avatar_component_clause, [], [f2448])).
fof(f2448, plain, (spl280_128 <=> (e12 = unit1)), introduced(avatar_definition, [new_symbols(naming, [spl280_128])])).
fof(f7640, plain, (~ spl280_181 | ~ spl280_182), inference(avatar_contradiction_clause, [], [f7639])).
fof(f7639, plain, ($false | (~ spl280_181 | ~ spl280_182)), inference(subsumption_resolution, [], [f7638, f969])).
fof(f7638, plain, ((e20 = e21) | (~ spl280_181 | ~ spl280_182)), inference(backward_demodulation, [], [f2727, f2723])).
fof(f2723, plain, ((e20 = op2(e22, e24)) | ~ spl280_181), inference(avatar_component_clause, [], [f2721])).
fof(f2721, plain, (spl280_181 <=> (e20 = op2(e22, e24))), introduced(avatar_definition, [new_symbols(naming, [spl280_181])])).
fof(f7607, plain, (~ spl280_237 | ~ spl280_239), inference(avatar_contradiction_clause, [], [f7606])).
fof(f7606, plain, ($false | (~ spl280_237 | ~ spl280_239)), inference(subsumption_resolution, [], [f7605, f974])).
fof(f7605, plain, ((e21 = e23) | (~ spl280_237 | ~ spl280_239)), inference(backward_demodulation, [], [f2966, f2958])).
fof(f2958, plain, ((e21 = op2(e20, e23)) | ~ spl280_237), inference(avatar_component_clause, [], [f2956])).
fof(f2956, plain, (spl280_237 <=> (e21 = op2(e20, e23))), introduced(avatar_definition, [new_symbols(naming, [spl280_237])])).
fof(f7587, plain, (spl280_189 | ~ spl280_258), inference(avatar_split_clause, [], [f7577, f3044, f2754])).
fof(f2754, plain, (spl280_189 <=> (e23 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl280_189])])).
fof(f3044, plain, (spl280_258 <=> (e22 = unit2)), introduced(avatar_definition, [new_symbols(naming, [spl280_258])])).
fof(f7577, plain, ((e23 = op2(e22, e23)) | ~ spl280_258), inference(backward_demodulation, [], [f704, f3046])).
fof(f3046, plain, ((e22 = unit2) | ~ spl280_258), inference(avatar_component_clause, [], [f3044])).
fof(f7570, plain, (~ spl280_250 | ~ spl280_235), inference(avatar_split_clause, [], [f7569, f2947, f3010])).
fof(f3010, plain, (spl280_250 <=> (e24 = op2(e20, e21))), introduced(avatar_definition, [new_symbols(naming, [spl280_250])])).
fof(f7569, plain, (~ (e24 = op2(e20, e21)) | ~ spl280_235), inference(backward_demodulation, [], [f915, f2949])).
fof(f915, plain, ~ (op2(e20, e21) = op2(e20, e24)), inference(cnf_transformation, [], [f8])).
fof(f7537, plain, (spl280_167 | ~ spl280_562), inference(avatar_split_clause, [], [f7188, f5125, f2662])).
fof(f7188, plain, ((e21 = op2(e23, e22)) | ~ spl280_562), inference(backward_demodulation, [], [f4887, f5126])).
fof(f7533, plain, (~ spl280_155 | ~ spl280_150), inference(avatar_split_clause, [], [f7530, f2590, f2611])).
fof(f2590, plain, (spl280_150 <=> (e24 = op2(e24, e21))), introduced(avatar_definition, [new_symbols(naming, [spl280_150])])).
fof(f7530, plain, (~ (e24 = op2(e24, e20)) | ~ spl280_150), inference(backward_demodulation, [], [f949, f2592])).
fof(f2592, plain, ((e24 = op2(e24, e21)) | ~ spl280_150), inference(avatar_component_clause, [], [f2590])).
fof(f949, plain, ~ (op2(e24, e20) = op2(e24, e21)), inference(cnf_transformation, [], [f8])).
fof(f7490, plain, (~ spl280_85 | ~ spl280_35), inference(avatar_split_clause, [], [f7489, f2057, f2267])).
fof(f2267, plain, (spl280_85 <=> (e14 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl280_85])])).
fof(f7489, plain, (~ (e14 = op1(e11, e13)) | ~ spl280_35), inference(forward_demodulation, [], [f794, f2059])).
fof(f794, plain, ~ (op1(e11, e13) = op1(e13, e13)), inference(cnf_transformation, [], [f7])).
fof(f7452, plain, (~ spl280_9 | ~ spl280_109), inference(avatar_split_clause, [], [f7451, f2368, f1948])).
fof(f1948, plain, (spl280_9 <=> (e13 = op1(e14, e13))), introduced(avatar_definition, [new_symbols(naming, [spl280_9])])).
fof(f7451, plain, (~ (e13 = op1(e14, e13)) | ~ spl280_109), inference(forward_demodulation, [], [f792, f2370])).
fof(f792, plain, ~ (op1(e10, e13) = op1(e14, e13)), inference(cnf_transformation, [], [f7])).
fof(f7444, plain, (~ spl280_11 | ~ spl280_13), inference(avatar_contradiction_clause, [], [f7443])).
fof(f7443, plain, ($false | (~ spl280_11 | ~ spl280_13)), inference(subsumption_resolution, [], [f7442, f960])).
fof(f7442, plain, ((e10 = e12) | (~ spl280_11 | ~ spl280_13)), inference(backward_demodulation, [], [f1967, f1959])).
fof(f1967, plain, ((e12 = op1(e14, e12)) | ~ spl280_13), inference(avatar_component_clause, [], [f1965])).
fof(f7373, plain, (~ spl280_84 | ~ spl280_109), inference(avatar_split_clause, [], [f7370, f2368, f2263])).
fof(f2263, plain, (spl280_84 <=> (e13 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl280_84])])).
fof(f7370, plain, (~ (e13 = op1(e11, e13)) | ~ spl280_109), inference(backward_demodulation, [], [f789, f2370])).
fof(f789, plain, ~ (op1(e10, e13) = op1(e11, e13)), inference(cnf_transformation, [], [f7])).
fof(f7344, plain, (~ spl280_91 | ~ spl280_127), inference(avatar_contradiction_clause, [], [f7343])).
fof(f7343, plain, ($false | (~ spl280_91 | ~ spl280_127)), inference(subsumption_resolution, [], [f7342, f959])).
fof(f7342, plain, ((e10 = e11) | (~ spl280_91 | ~ spl280_127)), inference(forward_demodulation, [], [f7327, f2295])).
fof(f7327, plain, ((e11 = op1(e11, e11)) | ~ spl280_127), inference(backward_demodulation, [], [f615, f2446])).
fof(f2446, plain, ((e11 = unit1) | ~ spl280_127), inference(avatar_component_clause, [], [f2444])).
fof(f2444, plain, (spl280_127 <=> (e11 = unit1)), introduced(avatar_definition, [new_symbols(naming, [spl280_127])])).
fof(f7305, plain, (spl280_553 | ~ spl280_162), inference(avatar_split_clause, [], [f7304, f2641, f5079])).
fof(f2641, plain, (spl280_162 <=> (e21 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl280_162])])).
fof(f7304, plain, ((e21 = h4(e13)) | ~ spl280_162), inference(backward_demodulation, [], [f1782, f2643])).
fof(f2643, plain, ((e21 = op2(e23, e23)) | ~ spl280_162), inference(avatar_component_clause, [], [f2641])).
fof(f7277, plain, (~ spl280_583 | ~ spl280_197 | ~ spl280_572), inference(avatar_split_clause, [], [f7275, f5189, f2788, f5248])).
fof(f2788, plain, (spl280_197 <=> (e21 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl280_197])])).
fof(f7275, plain, (~ (e21 = h2(e11)) | (~ spl280_197 | ~ spl280_572)), inference(backward_demodulation, [], [f7181, f2790])).
fof(f2790, plain, ((e21 = op2(e22, e21)) | ~ spl280_197), inference(avatar_component_clause, [], [f2788])).
fof(f7181, plain, (~ (op2(e22, e21) = h2(e11)) | ~ spl280_572), inference(backward_demodulation, [], [f870, f7169])).
fof(f7169, plain, ((op2(e20, e21) = h2(e11)) | ~ spl280_572), inference(backward_demodulation, [], [f4875, f5190])).
fof(f4875, plain, (h2(e11) = op2(h2(e13), e21)), inference(forward_demodulation, [], [f1771, f1772])).
fof(f1771, plain, (op2(op2(e21, e21), e21) = h2(e11)), inference(cnf_transformation, [], [f17])).
fof(f870, plain, ~ (op2(e20, e21) = op2(e22, e21)), inference(cnf_transformation, [], [f8])).
fof(f7233, plain, (spl280_583 | ~ spl280_247 | ~ spl280_572), inference(avatar_split_clause, [], [f7231, f5189, f2998, f5248])).
fof(f7231, plain, ((e21 = h2(e11)) | (~ spl280_247 | ~ spl280_572)), inference(backward_demodulation, [], [f7169, f3000])).
fof(f3000, plain, ((e21 = op2(e20, e21)) | ~ spl280_247), inference(avatar_component_clause, [], [f2998])).
fof(f7215, plain, (~ spl280_221 | ~ spl280_257), inference(avatar_contradiction_clause, [], [f7214])).
fof(f7214, plain, ($false | (~ spl280_221 | ~ spl280_257)), inference(subsumption_resolution, [], [f7213, f969])).
fof(f7213, plain, ((e20 = e21) | (~ spl280_221 | ~ spl280_257)), inference(forward_demodulation, [], [f7198, f2891])).
fof(f7198, plain, ((e21 = op2(e21, e21)) | ~ spl280_257), inference(backward_demodulation, [], [f701, f3042])).
fof(f3042, plain, ((e21 = unit2) | ~ spl280_257), inference(avatar_component_clause, [], [f3040])).
fof(f3040, plain, (spl280_257 <=> (e21 = unit2)), introduced(avatar_definition, [new_symbols(naming, [spl280_257])])).
fof(f7194, plain, (spl280_165 | ~ spl280_562), inference(avatar_split_clause, [], [f7189, f5125, f2653])).
fof(f7189, plain, ((e24 = op2(e23, e23)) | ~ spl280_562), inference(backward_demodulation, [], [f4888, f5126])).
fof(f4888, plain, (e24 = op2(h3(e13), h3(e13))), inference(backward_demodulation, [], [f1763, f1777])).
fof(f7191, plain, (~ spl280_189 | ~ spl280_562), inference(avatar_split_clause, [], [f7187, f5125, f2754])).
fof(f7187, plain, (~ (e23 = op2(e22, e23)) | ~ spl280_562), inference(backward_demodulation, [], [f4885, f5126])).
fof(f4885, plain, ~ (op2(e22, e23) = h3(e13)), inference(backward_demodulation, [], [f936, f1777])).
fof(f936, plain, ~ (op2(e22, e22) = op2(e22, e23)), inference(cnf_transformation, [], [f8])).
fof(f7190, plain, (~ spl280_219 | ~ spl280_562), inference(avatar_split_clause, [], [f7184, f5125, f2880])).
fof(f2880, plain, (spl280_219 <=> (e23 = op2(e21, e22))), introduced(avatar_definition, [new_symbols(naming, [spl280_219])])).
fof(f7184, plain, (~ (e23 = op2(e21, e22)) | ~ spl280_562), inference(backward_demodulation, [], [f4880, f5126])).
fof(f4880, plain, ~ (op2(e21, e22) = h3(e13)), inference(backward_demodulation, [], [f883, f1777])).
fof(f883, plain, ~ (op2(e21, e22) = op2(e22, e22)), inference(cnf_transformation, [], [f8])).
fof(f7177, plain, (~ spl280_211 | ~ spl280_572), inference(avatar_split_clause, [], [f7167, f5189, f2847])).
fof(f2847, plain, (spl280_211 <=> (e20 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl280_211])])).
fof(f7167, plain, (~ (e20 = op2(e21, e23)) | ~ spl280_572), inference(backward_demodulation, [], [f4872, f5190])).
fof(f4872, plain, ~ (op2(e21, e23) = h2(e13)), inference(backward_demodulation, [], [f924, f1772])).
fof(f924, plain, ~ (op2(e21, e21) = op2(e21, e23)), inference(cnf_transformation, [], [f8])).
fof(f7176, plain, (~ spl280_216 | ~ spl280_572), inference(avatar_split_clause, [], [f7166, f5189, f2868])).
fof(f2868, plain, (spl280_216 <=> (e20 = op2(e21, e22))), introduced(avatar_definition, [new_symbols(naming, [spl280_216])])).
fof(f7166, plain, (~ (e20 = op2(e21, e22)) | ~ spl280_572), inference(backward_demodulation, [], [f4871, f5190])).
fof(f4871, plain, ~ (op2(e21, e22) = h2(e13)), inference(backward_demodulation, [], [f923, f1772])).
fof(f923, plain, ~ (op2(e21, e21) = op2(e21, e22)), inference(cnf_transformation, [], [f8])).
fof(f7174, plain, (~ spl280_171 | ~ spl280_572), inference(avatar_split_clause, [], [f7164, f5189, f2679])).
fof(f2679, plain, (spl280_171 <=> (e20 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl280_171])])).
fof(f7164, plain, (~ (e20 = op2(e23, e21)) | ~ spl280_572), inference(backward_demodulation, [], [f4868, f5190])).
fof(f4868, plain, ~ (op2(e23, e21) = h2(e13)), inference(backward_demodulation, [], [f874, f1772])).
fof(f874, plain, ~ (op2(e21, e21) = op2(e23, e21)), inference(cnf_transformation, [], [f8])).
fof(f7173, plain, (~ spl280_196 | ~ spl280_572), inference(avatar_split_clause, [], [f7163, f5189, f2784])).
fof(f2784, plain, (spl280_196 <=> (e20 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl280_196])])).
fof(f7163, plain, (~ (e20 = op2(e22, e21)) | ~ spl280_572), inference(backward_demodulation, [], [f4867, f5190])).
fof(f4867, plain, ~ (op2(e22, e21) = h2(e13)), inference(backward_demodulation, [], [f873, f1772])).
fof(f873, plain, ~ (op2(e21, e21) = op2(e22, e21)), inference(cnf_transformation, [], [f8])).
fof(f7151, plain, (~ spl280_233 | ~ spl280_235), inference(avatar_contradiction_clause, [], [f7150])).
fof(f7150, plain, ($false | (~ spl280_233 | ~ spl280_235)), inference(subsumption_resolution, [], [f7149, f977])).
fof(f977, plain, ~ (e22 = e24), inference(cnf_transformation, [], [f10])).
fof(f7149, plain, ((e22 = e24) | (~ spl280_233 | ~ spl280_235)), inference(forward_demodulation, [], [f2949, f2941])).
fof(f2941, plain, ((e22 = op2(e20, e24)) | ~ spl280_233), inference(avatar_component_clause, [], [f2939])).
fof(f2939, plain, (spl280_233 <=> (e22 = op2(e20, e24))), introduced(avatar_definition, [new_symbols(naming, [spl280_233])])).
fof(f7059, plain, (~ spl280_113 | ~ spl280_103), inference(avatar_split_clause, [], [f7058, f2343, f2385])).
fof(f2343, plain, (spl280_103 <=> (e12 = op1(e10, e14))), introduced(avatar_definition, [new_symbols(naming, [spl280_103])])).
fof(f7058, plain, (~ (e12 = op1(e10, e12)) | ~ spl280_103), inference(forward_demodulation, [], [f817, f2345])).
fof(f2345, plain, ((e12 = op1(e10, e14)) | ~ spl280_103), inference(avatar_component_clause, [], [f2343])).
fof(f817, plain, ~ (op1(e10, e12) = op1(e10, e14)), inference(cnf_transformation, [], [f7])).
fof(f7050, plain, (~ spl280_97 | ~ spl280_22), inference(avatar_split_clause, [], [f7049, f2003, f2318])).
fof(f2003, plain, (spl280_22 <=> (e11 = op1(e14, e10))), introduced(avatar_definition, [new_symbols(naming, [spl280_22])])).
fof(f7049, plain, (~ (e11 = op1(e11, e10)) | ~ spl280_22), inference(forward_demodulation, [], [f765, f2005])).
fof(f2005, plain, ((e11 = op1(e14, e10)) | ~ spl280_22), inference(avatar_component_clause, [], [f2003])).
fof(f7044, plain, (~ spl280_89 | ~ spl280_64), inference(avatar_split_clause, [], [f7043, f2179, f2284])).
fof(f2284, plain, (spl280_89 <=> (e13 = op1(e11, e12))), introduced(avatar_definition, [new_symbols(naming, [spl280_89])])).
fof(f7043, plain, (~ (e13 = op1(e11, e12)) | ~ spl280_64), inference(forward_demodulation, [], [f783, f2181])).
fof(f783, plain, ~ (op1(e11, e12) = op1(e12, e12)), inference(cnf_transformation, [], [f7])).
fof(f6916, plain, (~ spl280_26 | ~ spl280_27), inference(avatar_contradiction_clause, [], [f6915])).
fof(f6915, plain, ($false | (~ spl280_26 | ~ spl280_27)), inference(subsumption_resolution, [], [f6914, f959])).
fof(f6914, plain, ((e10 = e11) | (~ spl280_26 | ~ spl280_27)), inference(backward_demodulation, [], [f2026, f2022])).
fof(f2026, plain, ((e11 = op1(e13, e14)) | ~ spl280_27), inference(avatar_component_clause, [], [f2024])).
fof(f2024, plain, (spl280_27 <=> (e11 = op1(e13, e14))), introduced(avatar_definition, [new_symbols(naming, [spl280_27])])).
fof(f6872, plain, (~ spl280_52 | ~ spl280_53), inference(avatar_contradiction_clause, [], [f6871])).
fof(f6871, plain, ($false | (~ spl280_52 | ~ spl280_53)), inference(subsumption_resolution, [], [f6870, f963])).
fof(f6870, plain, ((e11 = e12) | (~ spl280_52 | ~ spl280_53)), inference(backward_demodulation, [], [f2135, f2131])).
fof(f2135, plain, ((e12 = op1(e12, e14)) | ~ spl280_53), inference(avatar_component_clause, [], [f2133])).
fof(f2133, plain, (spl280_53 <=> (e12 = op1(e12, e14))), introduced(avatar_definition, [new_symbols(naming, [spl280_53])])).
fof(f6855, plain, (spl280_35 | ~ spl280_64), inference(avatar_split_clause, [], [f6849, f2179, f2057])).
fof(f6849, plain, ((e14 = op1(e13, e13)) | ~ spl280_64), inference(backward_demodulation, [], [f1759, f2181])).
fof(f1759, plain, (e14 = op1(op1(e12, e12), op1(e12, e12))), inference(cnf_transformation, [], [f14])).
fof(f14, plain, ((e14 = op1(op1(e12, e12), op1(e12, e12))) & (e13 = op1(e12, e12)) & (e11 = op1(op1(e12, e12), e12)) & (e10 = op1(op1(op1(e12, e12), e12), op1(op1(e12, e12), e12)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG051+1.p', ax14)).
fof(f6854, plain, (spl280_37 | ~ spl280_64), inference(avatar_split_clause, [], [f6848, f2179, f2066])).
fof(f6848, plain, ((e11 = op1(e13, e12)) | ~ spl280_64), inference(backward_demodulation, [], [f1757, f2181])).
fof(f1757, plain, (e11 = op1(op1(e12, e12), e12)), inference(cnf_transformation, [], [f14])).
fof(f6851, plain, (~ spl280_14 | ~ spl280_64), inference(avatar_split_clause, [], [f6845, f2179, f1969])).
fof(f1969, plain, (spl280_14 <=> (e13 = op1(e14, e12))), introduced(avatar_definition, [new_symbols(naming, [spl280_14])])).
fof(f6845, plain, (~ (e13 = op1(e14, e12)) | ~ spl280_64), inference(backward_demodulation, [], [f787, f2181])).
fof(f787, plain, ~ (op1(e12, e12) = op1(e14, e12)), inference(cnf_transformation, [], [f7])).
fof(f6814, plain, (~ spl280_78 | ~ spl280_79), inference(avatar_contradiction_clause, [], [f6813])).
fof(f6813, plain, ($false | (~ spl280_78 | ~ spl280_79)), inference(subsumption_resolution, [], [f6812, f966])).
fof(f966, plain, ~ (e12 = e13), inference(cnf_transformation, [], [f9])).
fof(f6812, plain, ((e12 = e13) | (~ spl280_78 | ~ spl280_79)), inference(backward_demodulation, [], [f2244, f2240])).
fof(f2240, plain, ((e12 = op1(e11, e14)) | ~ spl280_78), inference(avatar_component_clause, [], [f2238])).
fof(f2238, plain, (spl280_78 <=> (e12 = op1(e11, e14))), introduced(avatar_definition, [new_symbols(naming, [spl280_78])])).
fof(f6783, plain, (~ spl280_81 | ~ spl280_91), inference(avatar_split_clause, [], [f6774, f2293, f2251])).
fof(f2251, plain, (spl280_81 <=> (e10 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl280_81])])).
fof(f6774, plain, (~ (e10 = op1(e11, e13)) | ~ spl280_91), inference(backward_demodulation, [], [f824, f2295])).
fof(f824, plain, ~ (op1(e11, e11) = op1(e11, e13)), inference(cnf_transformation, [], [f7])).
fof(f6782, plain, (~ spl280_86 | ~ spl280_91), inference(avatar_split_clause, [], [f6773, f2293, f2272])).
fof(f2272, plain, (spl280_86 <=> (e10 = op1(e11, e12))), introduced(avatar_definition, [new_symbols(naming, [spl280_86])])).
fof(f6773, plain, (~ (e10 = op1(e11, e12)) | ~ spl280_91), inference(backward_demodulation, [], [f823, f2295])).
fof(f823, plain, ~ (op1(e11, e11) = op1(e11, e12)), inference(cnf_transformation, [], [f7])).
fof(f6780, plain, (~ spl280_41 | ~ spl280_91), inference(avatar_split_clause, [], [f6771, f2293, f2083])).
fof(f2083, plain, (spl280_41 <=> (e10 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl280_41])])).
fof(f6771, plain, (~ (e10 = op1(e13, e11)) | ~ spl280_91), inference(backward_demodulation, [], [f774, f2295])).
fof(f774, plain, ~ (op1(e11, e11) = op1(e13, e11)), inference(cnf_transformation, [], [f7])).
fof(f6741, plain, (~ spl280_80 | ~ spl280_105), inference(avatar_split_clause, [], [f6737, f2351, f2246])).
fof(f2246, plain, (spl280_80 <=> (e14 = op1(e11, e14))), introduced(avatar_definition, [new_symbols(naming, [spl280_80])])).
fof(f6737, plain, (~ (e14 = op1(e11, e14)) | ~ spl280_105), inference(backward_demodulation, [], [f799, f2353])).
fof(f799, plain, ~ (op1(e10, e14) = op1(e11, e14)), inference(cnf_transformation, [], [f7])).
fof(f6725, plain, (~ spl280_101 | ~ spl280_111), inference(avatar_split_clause, [], [f6718, f2377, f2335])).
fof(f6718, plain, (~ (e10 = op1(e10, e14)) | ~ spl280_111), inference(backward_demodulation, [], [f817, f2379])).
fof(f2379, plain, ((e10 = op1(e10, e12)) | ~ spl280_111), inference(avatar_component_clause, [], [f2377])).
fof(f6697, plain, (~ spl280_106 | ~ spl280_121), inference(avatar_split_clause, [], [f6686, f2419, f2356])).
fof(f2356, plain, (spl280_106 <=> (e10 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl280_106])])).
fof(f6686, plain, (~ (e10 = op1(e10, e13)) | ~ spl280_121), inference(backward_demodulation, [], [f811, f2421])).
fof(f811, plain, ~ (op1(e10, e10) = op1(e10, e13)), inference(cnf_transformation, [], [f7])).
fof(f6678, plain, (spl280_105 | ~ spl280_126), inference(avatar_split_clause, [], [f6664, f2440, f2351])).
fof(f6664, plain, ((e14 = op1(e10, e14)) | ~ spl280_126), inference(backward_demodulation, [], [f620, f2442])).
fof(f620, plain, (e14 = op1(unit1, e14)), inference(cnf_transformation, [], [f2])).
fof(f6676, plain, (spl280_109 | ~ spl280_126), inference(avatar_split_clause, [], [f6662, f2440, f2368])).
fof(f6662, plain, ((e13 = op1(e10, e13)) | ~ spl280_126), inference(backward_demodulation, [], [f618, f2442])).
fof(f618, plain, (e13 = op1(unit1, e13)), inference(cnf_transformation, [], [f2])).
fof(f6570, plain, (spl280_558 | ~ spl280_161), inference(avatar_split_clause, [], [f6569, f2637, f5104])).
fof(f5104, plain, (spl280_558 <=> (e20 = h4(e13))), introduced(avatar_definition, [new_symbols(naming, [spl280_558])])).
fof(f2637, plain, (spl280_161 <=> (e20 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl280_161])])).
fof(f6569, plain, ((e20 = h4(e13)) | ~ spl280_161), inference(backward_demodulation, [], [f1782, f2639])).
fof(f2639, plain, ((e20 = op2(e23, e23)) | ~ spl280_161), inference(avatar_component_clause, [], [f2637])).
fof(f6531, plain, (~ spl280_182 | ~ spl280_183), inference(avatar_contradiction_clause, [], [f6530])).
fof(f6530, plain, ($false | (~ spl280_182 | ~ spl280_183)), inference(subsumption_resolution, [], [f6529, f973])).
fof(f6529, plain, ((e21 = e22) | (~ spl280_182 | ~ spl280_183)), inference(backward_demodulation, [], [f2731, f2727])).
fof(f2731, plain, ((e22 = op2(e22, e24)) | ~ spl280_183), inference(avatar_component_clause, [], [f2729])).
fof(f2729, plain, (spl280_183 <=> (e22 = op2(e22, e24))), introduced(avatar_definition, [new_symbols(naming, [spl280_183])])).
fof(f6518, plain, (~ spl280_558 | ~ spl280_186), inference(avatar_split_clause, [], [f6514, f2742, f5104])).
fof(f6514, plain, (~ (e20 = h4(e13)) | ~ spl280_186), inference(backward_demodulation, [], [f4896, f2744])).
fof(f4896, plain, ~ (op2(e22, e23) = h4(e13)), inference(backward_demodulation, [], [f896, f1782])).
fof(f896, plain, ~ (op2(e22, e23) = op2(e23, e23)), inference(cnf_transformation, [], [f8])).
fof(f6508, plain, (spl280_562 | ~ spl280_194), inference(avatar_split_clause, [], [f6507, f2775, f5125])).
fof(f6507, plain, ((e23 = h3(e13)) | ~ spl280_194), inference(backward_demodulation, [], [f1777, f2777])).
fof(f6491, plain, (~ spl280_606 | ~ spl280_201), inference(avatar_split_clause, [], [f6484, f2805, f5363])).
fof(f2805, plain, (spl280_201 <=> (e20 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl280_201])])).
fof(f6484, plain, (~ (e20 = h1(e13)) | ~ spl280_201), inference(backward_demodulation, [], [f4855, f2807])).
fof(f2807, plain, ((e20 = op2(e22, e20)) | ~ spl280_201), inference(avatar_component_clause, [], [f2805])).
fof(f4855, plain, ~ (op2(e22, e20) = h1(e13)), inference(backward_demodulation, [], [f860, f1767])).
fof(f860, plain, ~ (op2(e20, e20) = op2(e22, e20)), inference(cnf_transformation, [], [f8])).
fof(f6473, plain, (~ spl280_208 | ~ spl280_209), inference(avatar_contradiction_clause, [], [f6472])).
fof(f6472, plain, ($false | (~ spl280_208 | ~ spl280_209)), inference(subsumption_resolution, [], [f6471, f976])).
fof(f976, plain, ~ (e22 = e23), inference(cnf_transformation, [], [f10])).
fof(f6471, plain, ((e22 = e23) | (~ spl280_208 | ~ spl280_209)), inference(backward_demodulation, [], [f2840, f2836])).
fof(f2836, plain, ((e22 = op2(e21, e24)) | ~ spl280_208), inference(avatar_component_clause, [], [f2834])).
fof(f2834, plain, (spl280_208 <=> (e22 = op2(e21, e24))), introduced(avatar_definition, [new_symbols(naming, [spl280_208])])).
fof(f6434, plain, (spl280_572 | ~ spl280_221), inference(avatar_split_clause, [], [f6433, f2889, f5189])).
fof(f6433, plain, ((e20 = h2(e13)) | ~ spl280_221), inference(backward_demodulation, [], [f1772, f2891])).
fof(f6380, plain, (~ spl280_231 | ~ spl280_241), inference(avatar_split_clause, [], [f6372, f2973, f2931])).
fof(f6372, plain, (~ (e20 = op2(e20, e24)) | ~ spl280_241), inference(backward_demodulation, [], [f917, f2975])).
fof(f917, plain, ~ (op2(e20, e22) = op2(e20, e24)), inference(cnf_transformation, [], [f8])).
fof(f6351, plain, (spl280_606 | ~ spl280_251), inference(avatar_split_clause, [], [f6350, f3015, f5363])).
fof(f6350, plain, ((e20 = h1(e13)) | ~ spl280_251), inference(backward_demodulation, [], [f1767, f3017])).
fof(f3017, plain, ((e20 = op2(e20, e20)) | ~ spl280_251), inference(avatar_component_clause, [], [f3015])).
fof(f6348, plain, (spl280_155 | ~ spl280_256), inference(avatar_split_clause, [], [f6334, f3036, f2611])).
fof(f6334, plain, ((e24 = op2(e24, e20)) | ~ spl280_256), inference(backward_demodulation, [], [f707, f3038])).
fof(f707, plain, (e24 = op2(e24, unit2)), inference(cnf_transformation, [], [f5])).
fof(f6347, plain, (spl280_235 | ~ spl280_256), inference(avatar_split_clause, [], [f6333, f3036, f2947])).
fof(f6333, plain, ((e24 = op2(e20, e24)) | ~ spl280_256), inference(backward_demodulation, [], [f706, f3038])).
fof(f706, plain, (e24 = op2(unit2, e24)), inference(cnf_transformation, [], [f5])).
fof(f6346, plain, (spl280_179 | ~ spl280_256), inference(avatar_split_clause, [], [f6332, f3036, f2712])).
fof(f6332, plain, ((e23 = op2(e23, e20)) | ~ spl280_256), inference(backward_demodulation, [], [f705, f3038])).
fof(f705, plain, (e23 = op2(e23, unit2)), inference(cnf_transformation, [], [f5])).
fof(f5883, plain, (spl280_570 | spl280_567 | spl280_564 | spl280_561 | ~ spl280_667 | ~ spl280_668 | ~ spl280_669 | ~ spl280_670 | ~ spl280_671 | ~ spl280_672 | ~ spl280_673 | ~ spl280_674 | ~ spl280_675 | ~ spl280_676 | ~ spl280_677 | ~ spl280_678 | ~ spl280_679 | ~ spl280_680 | ~ spl280_681 | ~ spl280_682 | ~ spl280_683 | ~ spl280_684 | ~ spl280_685 | ~ spl280_686 | ~ spl280_687 | ~ spl280_688 | ~ spl280_689 | ~ spl280_690 | ~ spl280_691), inference(avatar_split_clause, [], [f5782, f5880, f5876, f5872, f5868, f5864, f5860, f5856, f5852, f5848, f5844, f5840, f5836, f5832, f5828, f5824, f5820, f5816, f5812, f5808, f5804, f5800, f5796, f5792, f5788, f5784, f5121, f5139, f5158, f5177])).
fof(f5177, plain, (spl280_570 <=> sP268), introduced(avatar_definition, [new_symbols(naming, [spl280_570])])).
fof(f5158, plain, (spl280_567 <=> sP269), introduced(avatar_definition, [new_symbols(naming, [spl280_567])])).
fof(f5139, plain, (spl280_564 <=> sP270), introduced(avatar_definition, [new_symbols(naming, [spl280_564])])).
fof(f5121, plain, (spl280_561 <=> sP271), introduced(avatar_definition, [new_symbols(naming, [spl280_561])])).
fof(f5782, plain, (~ (h2(e14) = h3(op1(e10, e10))) | ~ (h2(e11) = h3(op1(e10, e11))) | ~ (h3(op1(e10, e12)) = op2(h2(e13), e22)) | ~ (h3(op1(e10, e13)) = op2(h2(e13), h3(e13))) | ~ (h3(op1(e10, e14)) = op2(h2(e13), e24)) | ~ (h3(op1(e11, e10)) = op2(e21, h2(e13))) | ~ (h2(e13) = h3(op1(e11, e11))) | ~ (op2(e21, e22) = h3(op1(e11, e12))) | ~ (h3(op1(e11, e13)) = op2(e21, h3(e13))) | ~ (op2(e21, e24) = h3(op1(e11, e14))) | ~ (h3(op1(e12, e10)) = op2(e22, h2(e13))) | ~ (op2(e22, e21) = h3(op1(e12, e11))) | ~ (h3(e13) = h3(op1(e12, e12))) | ~ (h3(op1(e12, e13)) = op2(e22, h3(e13))) | ~ (op2(e22, e24) = h3(op1(e12, e14))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h2(e13))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), e21)) | ~ (e21 = h3(op1(e13, e12))) | ~ (e24 = h3(op1(e13, e13))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), e24)) | ~ (h3(op1(e14, e10)) = op2(e24, h2(e13))) | ~ (op2(e24, e21) = h3(op1(e14, e11))) | ~ (op2(e24, e22) = h3(op1(e14, e12))) | ~ (h3(op1(e14, e13)) = op2(e24, h3(e13))) | ~ (h5(e13) = h3(op1(e14, e14))) | sP271 | sP270 | sP269 | sP268), inference(forward_demodulation, [], [f5781, f4874])).
fof(f5781, plain, (~ (op2(h2(e13), h2(e13)) = h3(op1(e10, e10))) | ~ (h2(e11) = h3(op1(e10, e11))) | ~ (h3(op1(e10, e12)) = op2(h2(e13), e22)) | ~ (h3(op1(e10, e13)) = op2(h2(e13), h3(e13))) | ~ (h3(op1(e10, e14)) = op2(h2(e13), e24)) | ~ (h3(op1(e11, e10)) = op2(e21, h2(e13))) | ~ (h2(e13) = h3(op1(e11, e11))) | ~ (op2(e21, e22) = h3(op1(e11, e12))) | ~ (h3(op1(e11, e13)) = op2(e21, h3(e13))) | ~ (op2(e21, e24) = h3(op1(e11, e14))) | ~ (h3(op1(e12, e10)) = op2(e22, h2(e13))) | ~ (op2(e22, e21) = h3(op1(e12, e11))) | ~ (h3(e13) = h3(op1(e12, e12))) | ~ (h3(op1(e12, e13)) = op2(e22, h3(e13))) | ~ (op2(e22, e24) = h3(op1(e12, e14))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h2(e13))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), e21)) | ~ (e21 = h3(op1(e13, e12))) | ~ (e24 = h3(op1(e13, e13))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), e24)) | ~ (h3(op1(e14, e10)) = op2(e24, h2(e13))) | ~ (op2(e24, e21) = h3(op1(e14, e11))) | ~ (op2(e24, e22) = h3(op1(e14, e12))) | ~ (h3(op1(e14, e13)) = op2(e24, h3(e13))) | ~ (h5(e13) = h3(op1(e14, e14))) | sP271 | sP270 | sP269 | sP268), inference(forward_demodulation, [], [f5780, f4893])).
fof(f5780, plain, (~ (h2(e11) = h3(op1(e10, e11))) | ~ (h3(op1(e10, e12)) = op2(h2(e13), e22)) | ~ (h3(op1(e10, e13)) = op2(h2(e13), h3(e13))) | ~ (h3(op1(e10, e14)) = op2(h2(e13), e24)) | ~ (h3(op1(e11, e10)) = op2(e21, h2(e13))) | ~ (h2(e13) = h3(op1(e11, e11))) | ~ (op2(e21, e22) = h3(op1(e11, e12))) | ~ (h3(op1(e11, e13)) = op2(e21, h3(e13))) | ~ (op2(e21, e24) = h3(op1(e11, e14))) | ~ (h3(op1(e12, e10)) = op2(e22, h2(e13))) | ~ (op2(e22, e21) = h3(op1(e12, e11))) | ~ (h3(e13) = h3(op1(e12, e12))) | ~ (h3(op1(e12, e13)) = op2(e22, h3(e13))) | ~ (op2(e22, e24) = h3(op1(e12, e14))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h2(e13))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), e21)) | ~ (e21 = h3(op1(e13, e12))) | ~ (e24 = h3(op1(e13, e13))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), e24)) | ~ (h3(op1(e14, e10)) = op2(e24, h2(e13))) | ~ (op2(e24, e21) = h3(op1(e14, e11))) | ~ (op2(e24, e22) = h3(op1(e14, e12))) | ~ (h3(op1(e14, e13)) = op2(e24, h3(e13))) | ~ (h5(e13) = h3(op1(e14, e14))) | sP271 | sP270 | sP269 | sP268 | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f5779, f4875])).
fof(f5779, plain, (~ (h3(op1(e10, e11)) = op2(h2(e13), e21)) | ~ (h3(op1(e10, e12)) = op2(h2(e13), e22)) | ~ (h3(op1(e10, e13)) = op2(h2(e13), h3(e13))) | ~ (h3(op1(e10, e14)) = op2(h2(e13), e24)) | ~ (h3(op1(e11, e10)) = op2(e21, h2(e13))) | ~ (h2(e13) = h3(op1(e11, e11))) | ~ (op2(e21, e22) = h3(op1(e11, e12))) | ~ (h3(op1(e11, e13)) = op2(e21, h3(e13))) | ~ (op2(e21, e24) = h3(op1(e11, e14))) | ~ (h3(op1(e12, e10)) = op2(e22, h2(e13))) | ~ (op2(e22, e21) = h3(op1(e12, e11))) | ~ (h3(e13) = h3(op1(e12, e12))) | ~ (h3(op1(e12, e13)) = op2(e22, h3(e13))) | ~ (op2(e22, e24) = h3(op1(e12, e14))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h2(e13))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), e21)) | ~ (e21 = h3(op1(e13, e12))) | ~ (e24 = h3(op1(e13, e13))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), e24)) | ~ (h3(op1(e14, e10)) = op2(e24, h2(e13))) | ~ (op2(e24, e21) = h3(op1(e14, e11))) | ~ (op2(e24, e22) = h3(op1(e14, e12))) | ~ (h3(op1(e14, e13)) = op2(e24, h3(e13))) | ~ (h5(e13) = h3(op1(e14, e14))) | sP271 | sP270 | sP269 | sP268 | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f5778, f4893])).
fof(f5778, plain, (~ (h3(op1(e10, e11)) = op2(h3(e10), e21)) | ~ (h3(op1(e10, e12)) = op2(h2(e13), e22)) | ~ (h3(op1(e10, e13)) = op2(h2(e13), h3(e13))) | ~ (h3(op1(e10, e14)) = op2(h2(e13), e24)) | ~ (h3(op1(e11, e10)) = op2(e21, h2(e13))) | ~ (h2(e13) = h3(op1(e11, e11))) | ~ (op2(e21, e22) = h3(op1(e11, e12))) | ~ (h3(op1(e11, e13)) = op2(e21, h3(e13))) | ~ (op2(e21, e24) = h3(op1(e11, e14))) | ~ (h3(op1(e12, e10)) = op2(e22, h2(e13))) | ~ (op2(e22, e21) = h3(op1(e12, e11))) | ~ (h3(e13) = h3(op1(e12, e12))) | ~ (h3(op1(e12, e13)) = op2(e22, h3(e13))) | ~ (op2(e22, e24) = h3(op1(e12, e14))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h2(e13))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), e21)) | ~ (e21 = h3(op1(e13, e12))) | ~ (e24 = h3(op1(e13, e13))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), e24)) | ~ (h3(op1(e14, e10)) = op2(e24, h2(e13))) | ~ (op2(e24, e21) = h3(op1(e14, e11))) | ~ (op2(e24, e22) = h3(op1(e14, e12))) | ~ (h3(op1(e14, e13)) = op2(e24, h3(e13))) | ~ (h5(e13) = h3(op1(e14, e14))) | sP271 | sP270 | sP269 | sP268 | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f5777, f4890])).
fof(f5777, plain, (~ (h3(op1(e10, e12)) = op2(h2(e13), e22)) | ~ (h3(op1(e10, e13)) = op2(h2(e13), h3(e13))) | ~ (h3(op1(e10, e14)) = op2(h2(e13), e24)) | ~ (h3(op1(e11, e10)) = op2(e21, h2(e13))) | ~ (h2(e13) = h3(op1(e11, e11))) | ~ (op2(e21, e22) = h3(op1(e11, e12))) | ~ (h3(op1(e11, e13)) = op2(e21, h3(e13))) | ~ (op2(e21, e24) = h3(op1(e11, e14))) | ~ (h3(op1(e12, e10)) = op2(e22, h2(e13))) | ~ (op2(e22, e21) = h3(op1(e12, e11))) | ~ (h3(e13) = h3(op1(e12, e12))) | ~ (h3(op1(e12, e13)) = op2(e22, h3(e13))) | ~ (op2(e22, e24) = h3(op1(e12, e14))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h2(e13))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), e21)) | ~ (e21 = h3(op1(e13, e12))) | ~ (e24 = h3(op1(e13, e13))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), e24)) | ~ (h3(op1(e14, e10)) = op2(e24, h2(e13))) | ~ (op2(e24, e21) = h3(op1(e14, e11))) | ~ (op2(e24, e22) = h3(op1(e14, e12))) | ~ (h3(op1(e14, e13)) = op2(e24, h3(e13))) | ~ (h5(e13) = h3(op1(e14, e14))) | sP271 | sP270 | sP269 | sP268 | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f5776, f4893])).
fof(f5776, plain, (~ (h3(op1(e10, e12)) = op2(h3(e10), e22)) | ~ (h3(op1(e10, e13)) = op2(h2(e13), h3(e13))) | ~ (h3(op1(e10, e14)) = op2(h2(e13), e24)) | ~ (h3(op1(e11, e10)) = op2(e21, h2(e13))) | ~ (h2(e13) = h3(op1(e11, e11))) | ~ (op2(e21, e22) = h3(op1(e11, e12))) | ~ (h3(op1(e11, e13)) = op2(e21, h3(e13))) | ~ (op2(e21, e24) = h3(op1(e11, e14))) | ~ (h3(op1(e12, e10)) = op2(e22, h2(e13))) | ~ (op2(e22, e21) = h3(op1(e12, e11))) | ~ (h3(e13) = h3(op1(e12, e12))) | ~ (h3(op1(e12, e13)) = op2(e22, h3(e13))) | ~ (op2(e22, e24) = h3(op1(e12, e14))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h2(e13))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), e21)) | ~ (e21 = h3(op1(e13, e12))) | ~ (e24 = h3(op1(e13, e13))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), e24)) | ~ (h3(op1(e14, e10)) = op2(e24, h2(e13))) | ~ (op2(e24, e21) = h3(op1(e14, e11))) | ~ (op2(e24, e22) = h3(op1(e14, e12))) | ~ (h3(op1(e14, e13)) = op2(e24, h3(e13))) | ~ (h5(e13) = h3(op1(e14, e14))) | sP271 | sP270 | sP269 | sP268 | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f5775, f1774])).
fof(f5775, plain, (~ (h3(op1(e10, e13)) = op2(h2(e13), h3(e13))) | ~ (h3(op1(e10, e14)) = op2(h2(e13), e24)) | ~ (h3(op1(e11, e10)) = op2(e21, h2(e13))) | ~ (h2(e13) = h3(op1(e11, e11))) | ~ (op2(e21, e22) = h3(op1(e11, e12))) | ~ (h3(op1(e11, e13)) = op2(e21, h3(e13))) | ~ (op2(e21, e24) = h3(op1(e11, e14))) | ~ (h3(op1(e12, e10)) = op2(e22, h2(e13))) | ~ (op2(e22, e21) = h3(op1(e12, e11))) | ~ (h3(e13) = h3(op1(e12, e12))) | ~ (h3(op1(e12, e13)) = op2(e22, h3(e13))) | ~ (op2(e22, e24) = h3(op1(e12, e14))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h2(e13))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), e21)) | ~ (e21 = h3(op1(e13, e12))) | ~ (e24 = h3(op1(e13, e13))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), e24)) | ~ (h3(op1(e14, e10)) = op2(e24, h2(e13))) | ~ (op2(e24, e21) = h3(op1(e14, e11))) | ~ (op2(e24, e22) = h3(op1(e14, e12))) | ~ (h3(op1(e14, e13)) = op2(e24, h3(e13))) | ~ (h5(e13) = h3(op1(e14, e14))) | sP271 | sP270 | sP269 | sP268 | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f5774, f4893])).
fof(f5774, plain, (~ (h3(op1(e10, e14)) = op2(h2(e13), e24)) | ~ (h3(op1(e11, e10)) = op2(e21, h2(e13))) | ~ (h2(e13) = h3(op1(e11, e11))) | ~ (op2(e21, e22) = h3(op1(e11, e12))) | ~ (h3(op1(e11, e13)) = op2(e21, h3(e13))) | ~ (op2(e21, e24) = h3(op1(e11, e14))) | ~ (h3(op1(e12, e10)) = op2(e22, h2(e13))) | ~ (op2(e22, e21) = h3(op1(e12, e11))) | ~ (h3(e13) = h3(op1(e12, e12))) | ~ (h3(op1(e12, e13)) = op2(e22, h3(e13))) | ~ (op2(e22, e24) = h3(op1(e12, e14))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h2(e13))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), e21)) | ~ (e21 = h3(op1(e13, e12))) | ~ (e24 = h3(op1(e13, e13))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), e24)) | ~ (h3(op1(e14, e10)) = op2(e24, h2(e13))) | ~ (op2(e24, e21) = h3(op1(e14, e11))) | ~ (op2(e24, e22) = h3(op1(e14, e12))) | ~ (h3(op1(e14, e13)) = op2(e24, h3(e13))) | ~ (h5(e13) = h3(op1(e14, e14))) | sP271 | sP270 | sP269 | sP268 | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f5773, f4893])).
fof(f5773, plain, (~ (h3(op1(e10, e14)) = op2(h3(e10), e24)) | ~ (h3(op1(e11, e10)) = op2(e21, h2(e13))) | ~ (h2(e13) = h3(op1(e11, e11))) | ~ (op2(e21, e22) = h3(op1(e11, e12))) | ~ (h3(op1(e11, e13)) = op2(e21, h3(e13))) | ~ (op2(e21, e24) = h3(op1(e11, e14))) | ~ (h3(op1(e12, e10)) = op2(e22, h2(e13))) | ~ (op2(e22, e21) = h3(op1(e12, e11))) | ~ (h3(e13) = h3(op1(e12, e12))) | ~ (h3(op1(e12, e13)) = op2(e22, h3(e13))) | ~ (op2(e22, e24) = h3(op1(e12, e14))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h2(e13))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), e21)) | ~ (e21 = h3(op1(e13, e12))) | ~ (e24 = h3(op1(e13, e13))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), e24)) | ~ (h3(op1(e14, e10)) = op2(e24, h2(e13))) | ~ (op2(e24, e21) = h3(op1(e14, e11))) | ~ (op2(e24, e22) = h3(op1(e14, e12))) | ~ (h3(op1(e14, e13)) = op2(e24, h3(e13))) | ~ (h5(e13) = h3(op1(e14, e14))) | sP271 | sP270 | sP269 | sP268 | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f5772, f4878])).
fof(f5772, plain, (~ (h3(op1(e11, e10)) = op2(e21, h2(e13))) | ~ (h2(e13) = h3(op1(e11, e11))) | ~ (op2(e21, e22) = h3(op1(e11, e12))) | ~ (h3(op1(e11, e13)) = op2(e21, h3(e13))) | ~ (op2(e21, e24) = h3(op1(e11, e14))) | ~ (h3(op1(e12, e10)) = op2(e22, h2(e13))) | ~ (op2(e22, e21) = h3(op1(e12, e11))) | ~ (h3(e13) = h3(op1(e12, e12))) | ~ (h3(op1(e12, e13)) = op2(e22, h3(e13))) | ~ (op2(e22, e24) = h3(op1(e12, e14))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h2(e13))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), e21)) | ~ (e21 = h3(op1(e13, e12))) | ~ (e24 = h3(op1(e13, e13))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), e24)) | ~ (h3(op1(e14, e10)) = op2(e24, h2(e13))) | ~ (op2(e24, e21) = h3(op1(e14, e11))) | ~ (op2(e24, e22) = h3(op1(e14, e12))) | ~ (h3(op1(e14, e13)) = op2(e24, h3(e13))) | ~ (h5(e13) = h3(op1(e14, e14))) | sP271 | sP270 | sP269 | sP268 | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f5771, f4890])).
fof(f5771, plain, (~ (h3(op1(e11, e10)) = op2(h3(e11), h2(e13))) | ~ (h2(e13) = h3(op1(e11, e11))) | ~ (op2(e21, e22) = h3(op1(e11, e12))) | ~ (h3(op1(e11, e13)) = op2(e21, h3(e13))) | ~ (op2(e21, e24) = h3(op1(e11, e14))) | ~ (h3(op1(e12, e10)) = op2(e22, h2(e13))) | ~ (op2(e22, e21) = h3(op1(e12, e11))) | ~ (h3(e13) = h3(op1(e12, e12))) | ~ (h3(op1(e12, e13)) = op2(e22, h3(e13))) | ~ (op2(e22, e24) = h3(op1(e12, e14))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h2(e13))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), e21)) | ~ (e21 = h3(op1(e13, e12))) | ~ (e24 = h3(op1(e13, e13))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), e24)) | ~ (h3(op1(e14, e10)) = op2(e24, h2(e13))) | ~ (op2(e24, e21) = h3(op1(e14, e11))) | ~ (op2(e24, e22) = h3(op1(e14, e12))) | ~ (h3(op1(e14, e13)) = op2(e24, h3(e13))) | ~ (h5(e13) = h3(op1(e14, e14))) | sP271 | sP270 | sP269 | sP268 | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f5770, f4893])).
fof(f5770, plain, (~ (h2(e13) = h3(op1(e11, e11))) | ~ (op2(e21, e22) = h3(op1(e11, e12))) | ~ (h3(op1(e11, e13)) = op2(e21, h3(e13))) | ~ (op2(e21, e24) = h3(op1(e11, e14))) | ~ (h3(op1(e12, e10)) = op2(e22, h2(e13))) | ~ (op2(e22, e21) = h3(op1(e12, e11))) | ~ (h3(e13) = h3(op1(e12, e12))) | ~ (h3(op1(e12, e13)) = op2(e22, h3(e13))) | ~ (op2(e22, e24) = h3(op1(e12, e14))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h2(e13))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), e21)) | ~ (e21 = h3(op1(e13, e12))) | ~ (e24 = h3(op1(e13, e13))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), e24)) | ~ (h3(op1(e14, e10)) = op2(e24, h2(e13))) | ~ (op2(e24, e21) = h3(op1(e14, e11))) | ~ (op2(e24, e22) = h3(op1(e14, e12))) | ~ (h3(op1(e14, e13)) = op2(e24, h3(e13))) | ~ (h5(e13) = h3(op1(e14, e14))) | sP271 | sP270 | sP269 | sP268 | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f5769, f1772])).
fof(f5769, plain, (~ (op2(e21, e21) = h3(op1(e11, e11))) | ~ (op2(e21, e22) = h3(op1(e11, e12))) | ~ (h3(op1(e11, e13)) = op2(e21, h3(e13))) | ~ (op2(e21, e24) = h3(op1(e11, e14))) | ~ (h3(op1(e12, e10)) = op2(e22, h2(e13))) | ~ (op2(e22, e21) = h3(op1(e12, e11))) | ~ (h3(e13) = h3(op1(e12, e12))) | ~ (h3(op1(e12, e13)) = op2(e22, h3(e13))) | ~ (op2(e22, e24) = h3(op1(e12, e14))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h2(e13))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), e21)) | ~ (e21 = h3(op1(e13, e12))) | ~ (e24 = h3(op1(e13, e13))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), e24)) | ~ (h3(op1(e14, e10)) = op2(e24, h2(e13))) | ~ (op2(e24, e21) = h3(op1(e14, e11))) | ~ (op2(e24, e22) = h3(op1(e14, e12))) | ~ (h3(op1(e14, e13)) = op2(e24, h3(e13))) | ~ (h5(e13) = h3(op1(e14, e14))) | sP271 | sP270 | sP269 | sP268 | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f5768, f4890])).
fof(f5768, plain, (~ (op2(e21, e22) = h3(op1(e11, e12))) | ~ (h3(op1(e11, e13)) = op2(e21, h3(e13))) | ~ (op2(e21, e24) = h3(op1(e11, e14))) | ~ (h3(op1(e12, e10)) = op2(e22, h2(e13))) | ~ (op2(e22, e21) = h3(op1(e12, e11))) | ~ (h3(e13) = h3(op1(e12, e12))) | ~ (h3(op1(e12, e13)) = op2(e22, h3(e13))) | ~ (op2(e22, e24) = h3(op1(e12, e14))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h2(e13))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), e21)) | ~ (e21 = h3(op1(e13, e12))) | ~ (e24 = h3(op1(e13, e13))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), e24)) | ~ (h3(op1(e14, e10)) = op2(e24, h2(e13))) | ~ (op2(e24, e21) = h3(op1(e14, e11))) | ~ (op2(e24, e22) = h3(op1(e14, e12))) | ~ (h3(op1(e14, e13)) = op2(e24, h3(e13))) | ~ (h5(e13) = h3(op1(e14, e14))) | sP271 | sP270 | sP269 | sP268 | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f5767, f4890])).
fof(f5767, plain, (~ (h3(op1(e11, e12)) = op2(h3(e11), e22)) | ~ (h3(op1(e11, e13)) = op2(e21, h3(e13))) | ~ (op2(e21, e24) = h3(op1(e11, e14))) | ~ (h3(op1(e12, e10)) = op2(e22, h2(e13))) | ~ (op2(e22, e21) = h3(op1(e12, e11))) | ~ (h3(e13) = h3(op1(e12, e12))) | ~ (h3(op1(e12, e13)) = op2(e22, h3(e13))) | ~ (op2(e22, e24) = h3(op1(e12, e14))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h2(e13))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), e21)) | ~ (e21 = h3(op1(e13, e12))) | ~ (e24 = h3(op1(e13, e13))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), e24)) | ~ (h3(op1(e14, e10)) = op2(e24, h2(e13))) | ~ (op2(e24, e21) = h3(op1(e14, e11))) | ~ (op2(e24, e22) = h3(op1(e14, e12))) | ~ (h3(op1(e14, e13)) = op2(e24, h3(e13))) | ~ (h5(e13) = h3(op1(e14, e14))) | sP271 | sP270 | sP269 | sP268 | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f5766, f1774])).
fof(f5766, plain, (~ (h3(op1(e11, e13)) = op2(e21, h3(e13))) | ~ (op2(e21, e24) = h3(op1(e11, e14))) | ~ (h3(op1(e12, e10)) = op2(e22, h2(e13))) | ~ (op2(e22, e21) = h3(op1(e12, e11))) | ~ (h3(e13) = h3(op1(e12, e12))) | ~ (h3(op1(e12, e13)) = op2(e22, h3(e13))) | ~ (op2(e22, e24) = h3(op1(e12, e14))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h2(e13))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), e21)) | ~ (e21 = h3(op1(e13, e12))) | ~ (e24 = h3(op1(e13, e13))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), e24)) | ~ (h3(op1(e14, e10)) = op2(e24, h2(e13))) | ~ (op2(e24, e21) = h3(op1(e14, e11))) | ~ (op2(e24, e22) = h3(op1(e14, e12))) | ~ (h3(op1(e14, e13)) = op2(e24, h3(e13))) | ~ (h5(e13) = h3(op1(e14, e14))) | sP271 | sP270 | sP269 | sP268 | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f5765, f4890])).
fof(f5765, plain, (~ (op2(e21, e24) = h3(op1(e11, e14))) | ~ (h3(op1(e12, e10)) = op2(e22, h2(e13))) | ~ (op2(e22, e21) = h3(op1(e12, e11))) | ~ (h3(e13) = h3(op1(e12, e12))) | ~ (h3(op1(e12, e13)) = op2(e22, h3(e13))) | ~ (op2(e22, e24) = h3(op1(e12, e14))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h2(e13))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), e21)) | ~ (e21 = h3(op1(e13, e12))) | ~ (e24 = h3(op1(e13, e13))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), e24)) | ~ (h3(op1(e14, e10)) = op2(e24, h2(e13))) | ~ (op2(e24, e21) = h3(op1(e14, e11))) | ~ (op2(e24, e22) = h3(op1(e14, e12))) | ~ (h3(op1(e14, e13)) = op2(e24, h3(e13))) | ~ (h5(e13) = h3(op1(e14, e14))) | sP271 | sP270 | sP269 | sP268 | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f5764, f4890])).
fof(f5764, plain, (~ (h3(op1(e11, e14)) = op2(h3(e11), e24)) | ~ (h3(op1(e12, e10)) = op2(e22, h2(e13))) | ~ (op2(e22, e21) = h3(op1(e12, e11))) | ~ (h3(e13) = h3(op1(e12, e12))) | ~ (h3(op1(e12, e13)) = op2(e22, h3(e13))) | ~ (op2(e22, e24) = h3(op1(e12, e14))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h2(e13))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), e21)) | ~ (e21 = h3(op1(e13, e12))) | ~ (e24 = h3(op1(e13, e13))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), e24)) | ~ (h3(op1(e14, e10)) = op2(e24, h2(e13))) | ~ (op2(e24, e21) = h3(op1(e14, e11))) | ~ (op2(e24, e22) = h3(op1(e14, e12))) | ~ (h3(op1(e14, e13)) = op2(e24, h3(e13))) | ~ (h5(e13) = h3(op1(e14, e14))) | sP271 | sP270 | sP269 | sP268 | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f5763, f4878])).
fof(f5763, plain, (~ (h3(op1(e12, e10)) = op2(e22, h2(e13))) | ~ (op2(e22, e21) = h3(op1(e12, e11))) | ~ (h3(e13) = h3(op1(e12, e12))) | ~ (h3(op1(e12, e13)) = op2(e22, h3(e13))) | ~ (op2(e22, e24) = h3(op1(e12, e14))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h2(e13))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), e21)) | ~ (e21 = h3(op1(e13, e12))) | ~ (e24 = h3(op1(e13, e13))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), e24)) | ~ (h3(op1(e14, e10)) = op2(e24, h2(e13))) | ~ (op2(e24, e21) = h3(op1(e14, e11))) | ~ (op2(e24, e22) = h3(op1(e14, e12))) | ~ (h3(op1(e14, e13)) = op2(e24, h3(e13))) | ~ (h5(e13) = h3(op1(e14, e14))) | sP271 | sP270 | sP269 | sP268 | ~ (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f5762, f1774])).
fof(f5762, plain, (~ (h3(op1(e12, e10)) = op2(h3(e12), h2(e13))) | ~ (op2(e22, e21) = h3(op1(e12, e11))) | ~ (h3(e13) = h3(op1(e12, e12))) | ~ (h3(op1(e12, e13)) = op2(e22, h3(e13))) | ~ (op2(e22, e24) = h3(op1(e12, e14))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h2(e13))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), e21)) | ~ (e21 = h3(op1(e13, e12))) | ~ (e24 = h3(op1(e13, e13))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), e24)) | ~ (h3(op1(e14, e10)) = op2(e24, h2(e13))) | ~ (op2(e24, e21) = h3(op1(e14, e11))) | ~ (op2(e24, e22) = h3(op1(e14, e12))) | ~ (h3(op1(e14, e13)) = op2(e24, h3(e13))) | ~ (h5(e13) = h3(op1(e14, e14))) | sP271 | sP270 | sP269 | sP268 | ~ (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f5761, f4893])).
fof(f5761, plain, (~ (op2(e22, e21) = h3(op1(e12, e11))) | ~ (h3(e13) = h3(op1(e12, e12))) | ~ (h3(op1(e12, e13)) = op2(e22, h3(e13))) | ~ (op2(e22, e24) = h3(op1(e12, e14))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h2(e13))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), e21)) | ~ (e21 = h3(op1(e13, e12))) | ~ (e24 = h3(op1(e13, e13))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), e24)) | ~ (h3(op1(e14, e10)) = op2(e24, h2(e13))) | ~ (op2(e24, e21) = h3(op1(e14, e11))) | ~ (op2(e24, e22) = h3(op1(e14, e12))) | ~ (h3(op1(e14, e13)) = op2(e24, h3(e13))) | ~ (h5(e13) = h3(op1(e14, e14))) | sP271 | sP270 | sP269 | sP268 | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f5760, f1774])).
fof(f5760, plain, (~ (h3(op1(e12, e11)) = op2(h3(e12), e21)) | ~ (h3(e13) = h3(op1(e12, e12))) | ~ (h3(op1(e12, e13)) = op2(e22, h3(e13))) | ~ (op2(e22, e24) = h3(op1(e12, e14))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h2(e13))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), e21)) | ~ (e21 = h3(op1(e13, e12))) | ~ (e24 = h3(op1(e13, e13))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), e24)) | ~ (h3(op1(e14, e10)) = op2(e24, h2(e13))) | ~ (op2(e24, e21) = h3(op1(e14, e11))) | ~ (op2(e24, e22) = h3(op1(e14, e12))) | ~ (h3(op1(e14, e13)) = op2(e24, h3(e13))) | ~ (h5(e13) = h3(op1(e14, e14))) | sP271 | sP270 | sP269 | sP268 | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f5759, f4890])).
fof(f5759, plain, (~ (h3(e13) = h3(op1(e12, e12))) | ~ (h3(op1(e12, e13)) = op2(e22, h3(e13))) | ~ (op2(e22, e24) = h3(op1(e12, e14))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h2(e13))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), e21)) | ~ (e21 = h3(op1(e13, e12))) | ~ (e24 = h3(op1(e13, e13))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), e24)) | ~ (h3(op1(e14, e10)) = op2(e24, h2(e13))) | ~ (op2(e24, e21) = h3(op1(e14, e11))) | ~ (op2(e24, e22) = h3(op1(e14, e12))) | ~ (h3(op1(e14, e13)) = op2(e24, h3(e13))) | ~ (h5(e13) = h3(op1(e14, e14))) | sP271 | sP270 | sP269 | sP268 | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f5758, f1777])).
fof(f5758, plain, (~ (op2(e22, e22) = h3(op1(e12, e12))) | ~ (h3(op1(e12, e13)) = op2(e22, h3(e13))) | ~ (op2(e22, e24) = h3(op1(e12, e14))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h2(e13))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), e21)) | ~ (e21 = h3(op1(e13, e12))) | ~ (e24 = h3(op1(e13, e13))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), e24)) | ~ (h3(op1(e14, e10)) = op2(e24, h2(e13))) | ~ (op2(e24, e21) = h3(op1(e14, e11))) | ~ (op2(e24, e22) = h3(op1(e14, e12))) | ~ (h3(op1(e14, e13)) = op2(e24, h3(e13))) | ~ (h5(e13) = h3(op1(e14, e14))) | sP271 | sP270 | sP269 | sP268 | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f5757, f1774])).
fof(f5757, plain, (~ (h3(op1(e12, e13)) = op2(e22, h3(e13))) | ~ (op2(e22, e24) = h3(op1(e12, e14))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h2(e13))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), e21)) | ~ (e21 = h3(op1(e13, e12))) | ~ (e24 = h3(op1(e13, e13))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), e24)) | ~ (h3(op1(e14, e10)) = op2(e24, h2(e13))) | ~ (op2(e24, e21) = h3(op1(e14, e11))) | ~ (op2(e24, e22) = h3(op1(e14, e12))) | ~ (h3(op1(e14, e13)) = op2(e24, h3(e13))) | ~ (h5(e13) = h3(op1(e14, e14))) | sP271 | sP270 | sP269 | sP268 | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f5756, f1774])).
fof(f5756, plain, (~ (op2(e22, e24) = h3(op1(e12, e14))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h2(e13))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), e21)) | ~ (e21 = h3(op1(e13, e12))) | ~ (e24 = h3(op1(e13, e13))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), e24)) | ~ (h3(op1(e14, e10)) = op2(e24, h2(e13))) | ~ (op2(e24, e21) = h3(op1(e14, e11))) | ~ (op2(e24, e22) = h3(op1(e14, e12))) | ~ (h3(op1(e14, e13)) = op2(e24, h3(e13))) | ~ (h5(e13) = h3(op1(e14, e14))) | sP271 | sP270 | sP269 | sP268 | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f5755, f1774])).
fof(f5755, plain, (~ (h3(op1(e12, e14)) = op2(h3(e12), e24)) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h2(e13))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), e21)) | ~ (e21 = h3(op1(e13, e12))) | ~ (e24 = h3(op1(e13, e13))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), e24)) | ~ (h3(op1(e14, e10)) = op2(e24, h2(e13))) | ~ (op2(e24, e21) = h3(op1(e14, e11))) | ~ (op2(e24, e22) = h3(op1(e14, e12))) | ~ (h3(op1(e14, e13)) = op2(e24, h3(e13))) | ~ (h5(e13) = h3(op1(e14, e14))) | sP271 | sP270 | sP269 | sP268 | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f5754, f4878])).
fof(f5754, plain, (~ (h3(op1(e13, e10)) = op2(h3(e13), h2(e13))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), e21)) | ~ (e21 = h3(op1(e13, e12))) | ~ (e24 = h3(op1(e13, e13))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), e24)) | ~ (h3(op1(e14, e10)) = op2(e24, h2(e13))) | ~ (op2(e24, e21) = h3(op1(e14, e11))) | ~ (op2(e24, e22) = h3(op1(e14, e12))) | ~ (h3(op1(e14, e13)) = op2(e24, h3(e13))) | ~ (h5(e13) = h3(op1(e14, e14))) | sP271 | sP270 | sP269 | sP268 | ~ (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f5753, f4893])).
fof(f5753, plain, (~ (h3(op1(e13, e11)) = op2(h3(e13), e21)) | ~ (e21 = h3(op1(e13, e12))) | ~ (e24 = h3(op1(e13, e13))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), e24)) | ~ (h3(op1(e14, e10)) = op2(e24, h2(e13))) | ~ (op2(e24, e21) = h3(op1(e14, e11))) | ~ (op2(e24, e22) = h3(op1(e14, e12))) | ~ (h3(op1(e14, e13)) = op2(e24, h3(e13))) | ~ (h5(e13) = h3(op1(e14, e14))) | sP271 | sP270 | sP269 | sP268 | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f5752, f4890])).
fof(f5752, plain, (~ (e21 = h3(op1(e13, e12))) | ~ (e24 = h3(op1(e13, e13))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), e24)) | ~ (h3(op1(e14, e10)) = op2(e24, h2(e13))) | ~ (op2(e24, e21) = h3(op1(e14, e11))) | ~ (op2(e24, e22) = h3(op1(e14, e12))) | ~ (h3(op1(e14, e13)) = op2(e24, h3(e13))) | ~ (h5(e13) = h3(op1(e14, e14))) | sP271 | sP270 | sP269 | sP268 | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f5751, f4887])).
fof(f5751, plain, (~ (h3(op1(e13, e12)) = op2(h3(e13), e22)) | ~ (e24 = h3(op1(e13, e13))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), e24)) | ~ (h3(op1(e14, e10)) = op2(e24, h2(e13))) | ~ (op2(e24, e21) = h3(op1(e14, e11))) | ~ (op2(e24, e22) = h3(op1(e14, e12))) | ~ (h3(op1(e14, e13)) = op2(e24, h3(e13))) | ~ (h5(e13) = h3(op1(e14, e14))) | sP271 | sP270 | sP269 | sP268 | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f5750, f1774])).
fof(f5750, plain, (~ (e24 = h3(op1(e13, e13))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), e24)) | ~ (h3(op1(e14, e10)) = op2(e24, h2(e13))) | ~ (op2(e24, e21) = h3(op1(e14, e11))) | ~ (op2(e24, e22) = h3(op1(e14, e12))) | ~ (h3(op1(e14, e13)) = op2(e24, h3(e13))) | ~ (h5(e13) = h3(op1(e14, e14))) | sP271 | sP270 | sP269 | sP268 | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f5749, f4888])).
fof(f5749, plain, (~ (h3(op1(e13, e14)) = op2(h3(e13), e24)) | ~ (h3(op1(e14, e10)) = op2(e24, h2(e13))) | ~ (op2(e24, e21) = h3(op1(e14, e11))) | ~ (op2(e24, e22) = h3(op1(e14, e12))) | ~ (h3(op1(e14, e13)) = op2(e24, h3(e13))) | ~ (h5(e13) = h3(op1(e14, e14))) | sP271 | sP270 | sP269 | sP268 | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f5748, f4878])).
fof(f5748, plain, (~ (h3(op1(e14, e10)) = op2(e24, h2(e13))) | ~ (op2(e24, e21) = h3(op1(e14, e11))) | ~ (op2(e24, e22) = h3(op1(e14, e12))) | ~ (h3(op1(e14, e13)) = op2(e24, h3(e13))) | ~ (h5(e13) = h3(op1(e14, e14))) | sP271 | sP270 | sP269 | sP268 | ~ (h3(op1(e13, e14)) = op2(h3(e13), h3(e14))) | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f5747, f4878])).
fof(f5747, plain, (~ (h3(op1(e14, e10)) = op2(h3(e14), h2(e13))) | ~ (op2(e24, e21) = h3(op1(e14, e11))) | ~ (op2(e24, e22) = h3(op1(e14, e12))) | ~ (h3(op1(e14, e13)) = op2(e24, h3(e13))) | ~ (h5(e13) = h3(op1(e14, e14))) | sP271 | sP270 | sP269 | sP268 | ~ (h3(op1(e13, e14)) = op2(h3(e13), h3(e14))) | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f5746, f4893])).
fof(f5746, plain, (~ (op2(e24, e21) = h3(op1(e14, e11))) | ~ (op2(e24, e22) = h3(op1(e14, e12))) | ~ (h3(op1(e14, e13)) = op2(e24, h3(e13))) | ~ (h5(e13) = h3(op1(e14, e14))) | sP271 | sP270 | sP269 | sP268 | ~ (h3(op1(e14, e10)) = op2(h3(e14), h3(e10))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), h3(e14))) | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f5745, f4878])).
fof(f5745, plain, (~ (h3(op1(e14, e11)) = op2(h3(e14), e21)) | ~ (op2(e24, e22) = h3(op1(e14, e12))) | ~ (h3(op1(e14, e13)) = op2(e24, h3(e13))) | ~ (h5(e13) = h3(op1(e14, e14))) | sP271 | sP270 | sP269 | sP268 | ~ (h3(op1(e14, e10)) = op2(h3(e14), h3(e10))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), h3(e14))) | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f5744, f4890])).
fof(f5744, plain, (~ (op2(e24, e22) = h3(op1(e14, e12))) | ~ (h3(op1(e14, e13)) = op2(e24, h3(e13))) | ~ (h5(e13) = h3(op1(e14, e14))) | sP271 | sP270 | sP269 | sP268 | ~ (h3(op1(e14, e11)) = op2(h3(e14), h3(e11))) | ~ (h3(op1(e14, e10)) = op2(h3(e14), h3(e10))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), h3(e14))) | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f5743, f4878])).
fof(f5743, plain, (~ (h3(op1(e14, e12)) = op2(h3(e14), e22)) | ~ (h3(op1(e14, e13)) = op2(e24, h3(e13))) | ~ (h5(e13) = h3(op1(e14, e14))) | sP271 | sP270 | sP269 | sP268 | ~ (h3(op1(e14, e11)) = op2(h3(e14), h3(e11))) | ~ (h3(op1(e14, e10)) = op2(h3(e14), h3(e10))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), h3(e14))) | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f5742, f1774])).
fof(f5742, plain, (~ (h3(op1(e14, e13)) = op2(e24, h3(e13))) | ~ (h5(e13) = h3(op1(e14, e14))) | sP271 | sP270 | sP269 | sP268 | ~ (h3(op1(e14, e12)) = op2(h3(e14), h3(e12))) | ~ (h3(op1(e14, e11)) = op2(h3(e14), h3(e11))) | ~ (h3(op1(e14, e10)) = op2(h3(e14), h3(e10))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), h3(e14))) | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f5741, f4878])).
fof(f5741, plain, (~ (h5(e13) = h3(op1(e14, e14))) | sP271 | sP270 | sP269 | sP268 | ~ (h3(op1(e14, e13)) = op2(h3(e14), h3(e13))) | ~ (h3(op1(e14, e12)) = op2(h3(e14), h3(e12))) | ~ (h3(op1(e14, e11)) = op2(h3(e14), h3(e11))) | ~ (h3(op1(e14, e10)) = op2(h3(e14), h3(e10))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), h3(e14))) | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f5740, f1787])).
fof(f5740, plain, (~ (op2(e24, e24) = h3(op1(e14, e14))) | sP271 | sP270 | sP269 | sP268 | ~ (h3(op1(e14, e13)) = op2(h3(e14), h3(e13))) | ~ (h3(op1(e14, e12)) = op2(h3(e14), h3(e12))) | ~ (h3(op1(e14, e11)) = op2(h3(e14), h3(e11))) | ~ (h3(op1(e14, e10)) = op2(h3(e14), h3(e10))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), h3(e14))) | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f5739, f4878])).
fof(f5739, plain, (sP271 | sP270 | sP269 | sP268 | ~ (h3(op1(e14, e14)) = op2(h3(e14), h3(e14))) | ~ (h3(op1(e14, e13)) = op2(h3(e14), h3(e13))) | ~ (h3(op1(e14, e12)) = op2(h3(e14), h3(e12))) | ~ (h3(op1(e14, e11)) = op2(h3(e14), h3(e11))) | ~ (h3(op1(e14, e10)) = op2(h3(e14), h3(e10))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), h3(e14))) | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(subsumption_resolution, [], [f1903, f4878])).
fof(f1903, plain, (~ (e24 = h3(e14)) | sP271 | sP270 | sP269 | sP268 | ~ (h3(op1(e14, e14)) = op2(h3(e14), h3(e14))) | ~ (h3(op1(e14, e13)) = op2(h3(e14), h3(e13))) | ~ (h3(op1(e14, e12)) = op2(h3(e14), h3(e12))) | ~ (h3(op1(e14, e11)) = op2(h3(e14), h3(e11))) | ~ (h3(op1(e14, e10)) = op2(h3(e14), h3(e10))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), h3(e14))) | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(cnf_transformation, [], [f306])).
fof(f306, plain, (((~ (e24 = h5(e14)) & ~ (e24 = h5(e13)) & ~ (e24 = h5(e12)) & ~ (e24 = h5(e11)) & ~ (e24 = h5(e10))) | sP279 | sP278 | sP277 | sP276 | ~ (h5(op1(e14, e14)) = op2(h5(e14), h5(e14))) | ~ (h5(op1(e14, e13)) = op2(h5(e14), h5(e13))) | ~ (h5(op1(e14, e12)) = op2(h5(e14), h5(e12))) | ~ (h5(op1(e14, e11)) = op2(h5(e14), h5(e11))) | ~ (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) | ~ (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))) & ((~ (e24 = h4(e14)) & ~ (e24 = h4(e13)) & ~ (e24 = h4(e12)) & ~ (e24 = h4(e11)) & ~ (e24 = h4(e10))) | sP275 | sP274 | sP273 | sP272 | ~ (h4(op1(e14, e14)) = op2(h4(e14), h4(e14))) | ~ (h4(op1(e14, e13)) = op2(h4(e14), h4(e13))) | ~ (h4(op1(e14, e12)) = op2(h4(e14), h4(e12))) | ~ (h4(op1(e14, e11)) = op2(h4(e14), h4(e11))) | ~ (h4(op1(e14, e10)) = op2(h4(e14), h4(e10))) | ~ (h4(op1(e13, e14)) = op2(h4(e13), h4(e14))) | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e14)) = op2(h4(e12), h4(e14))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e14)) = op2(h4(e11), h4(e14))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e14)) = op2(h4(e10), h4(e14))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) & ((~ (e24 = h3(e14)) & ~ (e24 = h3(e13)) & ~ (e24 = h3(e12)) & ~ (e24 = h3(e11)) & ~ (e24 = h3(e10))) | sP271 | sP270 | sP269 | sP268 | ~ (h3(op1(e14, e14)) = op2(h3(e14), h3(e14))) | ~ (h3(op1(e14, e13)) = op2(h3(e14), h3(e13))) | ~ (h3(op1(e14, e12)) = op2(h3(e14), h3(e12))) | ~ (h3(op1(e14, e11)) = op2(h3(e14), h3(e11))) | ~ (h3(op1(e14, e10)) = op2(h3(e14), h3(e10))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), h3(e14))) | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) & ((~ (e24 = h2(e14)) & ~ (e24 = h2(e13)) & ~ (e24 = h2(e12)) & ~ (e24 = h2(e11)) & ~ (e24 = h2(e10))) | sP267 | sP266 | sP265 | sP264 | ~ (h2(op1(e14, e14)) = op2(h2(e14), h2(e14))) | ~ (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) | ~ (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) | ~ (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) | ~ (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) | ~ (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) & ((~ (e24 = h1(e14)) & ~ (e24 = h1(e13)) & ~ (e24 = h1(e12)) & ~ (e24 = h1(e11)) & ~ (e24 = h1(e10))) | sP263 | sP262 | sP261 | sP260 | ~ (h1(op1(e14, e14)) = op2(h1(e14), h1(e14))) | ~ (h1(op1(e14, e13)) = op2(h1(e14), h1(e13))) | ~ (h1(op1(e14, e12)) = op2(h1(e14), h1(e12))) | ~ (h1(op1(e14, e11)) = op2(h1(e14), h1(e11))) | ~ (h1(op1(e14, e10)) = op2(h1(e14), h1(e10))) | ~ (h1(op1(e13, e14)) = op2(h1(e13), h1(e14))) | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e14)) = op2(h1(e12), h1(e14))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e14)) = op2(h1(e11), h1(e14))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e14)) = op2(h1(e10), h1(e14))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(definition_folding, [], [f23, e305, e304, e303, e302, e301, e300, e299, e298, e297, e296, e295, e294, e293, e292, e291, e290, e289, e288, e287, e286])).
fof(f286, plain, ((~ (e20 = h1(e14)) & ~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10))) | ~ sP260), inference(usedef, [], [e286])).
fof(e286, plain, (sP260 <=> (~ (e20 = h1(e14)) & ~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP260])])).
fof(f287, plain, ((~ (e21 = h1(e14)) & ~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10))) | ~ sP261), inference(usedef, [], [e287])).
fof(e287, plain, (sP261 <=> (~ (e21 = h1(e14)) & ~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP261])])).
fof(f288, plain, ((~ (e22 = h1(e14)) & ~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10))) | ~ sP262), inference(usedef, [], [e288])).
fof(e288, plain, (sP262 <=> (~ (e22 = h1(e14)) & ~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP262])])).
fof(f289, plain, ((~ (e23 = h1(e14)) & ~ (e23 = h1(e13)) & ~ (e23 = h1(e12)) & ~ (e23 = h1(e11)) & ~ (e23 = h1(e10))) | ~ sP263), inference(usedef, [], [e289])).
fof(e289, plain, (sP263 <=> (~ (e23 = h1(e14)) & ~ (e23 = h1(e13)) & ~ (e23 = h1(e12)) & ~ (e23 = h1(e11)) & ~ (e23 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP263])])).
fof(f290, plain, ((~ (e20 = h2(e14)) & ~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10))) | ~ sP264), inference(usedef, [], [e290])).
fof(e290, plain, (sP264 <=> (~ (e20 = h2(e14)) & ~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP264])])).
fof(f291, plain, ((~ (e21 = h2(e14)) & ~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10))) | ~ sP265), inference(usedef, [], [e291])).
fof(e291, plain, (sP265 <=> (~ (e21 = h2(e14)) & ~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP265])])).
fof(f292, plain, ((~ (e22 = h2(e14)) & ~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10))) | ~ sP266), inference(usedef, [], [e292])).
fof(e292, plain, (sP266 <=> (~ (e22 = h2(e14)) & ~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP266])])).
fof(f293, plain, ((~ (e23 = h2(e14)) & ~ (e23 = h2(e13)) & ~ (e23 = h2(e12)) & ~ (e23 = h2(e11)) & ~ (e23 = h2(e10))) | ~ sP267), inference(usedef, [], [e293])).
fof(e293, plain, (sP267 <=> (~ (e23 = h2(e14)) & ~ (e23 = h2(e13)) & ~ (e23 = h2(e12)) & ~ (e23 = h2(e11)) & ~ (e23 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP267])])).
fof(f294, plain, ((~ (e20 = h3(e14)) & ~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10))) | ~ sP268), inference(usedef, [], [e294])).
fof(e294, plain, (sP268 <=> (~ (e20 = h3(e14)) & ~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP268])])).
fof(f295, plain, ((~ (e21 = h3(e14)) & ~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10))) | ~ sP269), inference(usedef, [], [e295])).
fof(e295, plain, (sP269 <=> (~ (e21 = h3(e14)) & ~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP269])])).
fof(f296, plain, ((~ (e22 = h3(e14)) & ~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10))) | ~ sP270), inference(usedef, [], [e296])).
fof(e296, plain, (sP270 <=> (~ (e22 = h3(e14)) & ~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP270])])).
fof(f297, plain, ((~ (e23 = h3(e14)) & ~ (e23 = h3(e13)) & ~ (e23 = h3(e12)) & ~ (e23 = h3(e11)) & ~ (e23 = h3(e10))) | ~ sP271), inference(usedef, [], [e297])).
fof(e297, plain, (sP271 <=> (~ (e23 = h3(e14)) & ~ (e23 = h3(e13)) & ~ (e23 = h3(e12)) & ~ (e23 = h3(e11)) & ~ (e23 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP271])])).
fof(f298, plain, ((~ (e20 = h4(e14)) & ~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10))) | ~ sP272), inference(usedef, [], [e298])).
fof(e298, plain, (sP272 <=> (~ (e20 = h4(e14)) & ~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP272])])).
fof(f299, plain, ((~ (e21 = h4(e14)) & ~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10))) | ~ sP273), inference(usedef, [], [e299])).
fof(e299, plain, (sP273 <=> (~ (e21 = h4(e14)) & ~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP273])])).
fof(f300, plain, ((~ (e22 = h4(e14)) & ~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10))) | ~ sP274), inference(usedef, [], [e300])).
fof(e300, plain, (sP274 <=> (~ (e22 = h4(e14)) & ~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP274])])).
fof(f301, plain, ((~ (e23 = h4(e14)) & ~ (e23 = h4(e13)) & ~ (e23 = h4(e12)) & ~ (e23 = h4(e11)) & ~ (e23 = h4(e10))) | ~ sP275), inference(usedef, [], [e301])).
fof(e301, plain, (sP275 <=> (~ (e23 = h4(e14)) & ~ (e23 = h4(e13)) & ~ (e23 = h4(e12)) & ~ (e23 = h4(e11)) & ~ (e23 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP275])])).
fof(f302, plain, ((~ (e20 = h5(e14)) & ~ (e20 = h5(e13)) & ~ (e20 = h5(e12)) & ~ (e20 = h5(e11)) & ~ (e20 = h5(e10))) | ~ sP276), inference(usedef, [], [e302])).
fof(e302, plain, (sP276 <=> (~ (e20 = h5(e14)) & ~ (e20 = h5(e13)) & ~ (e20 = h5(e12)) & ~ (e20 = h5(e11)) & ~ (e20 = h5(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP276])])).
fof(f303, plain, ((~ (e21 = h5(e14)) & ~ (e21 = h5(e13)) & ~ (e21 = h5(e12)) & ~ (e21 = h5(e11)) & ~ (e21 = h5(e10))) | ~ sP277), inference(usedef, [], [e303])).
fof(e303, plain, (sP277 <=> (~ (e21 = h5(e14)) & ~ (e21 = h5(e13)) & ~ (e21 = h5(e12)) & ~ (e21 = h5(e11)) & ~ (e21 = h5(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP277])])).
fof(f304, plain, ((~ (e22 = h5(e14)) & ~ (e22 = h5(e13)) & ~ (e22 = h5(e12)) & ~ (e22 = h5(e11)) & ~ (e22 = h5(e10))) | ~ sP278), inference(usedef, [], [e304])).
fof(e304, plain, (sP278 <=> (~ (e22 = h5(e14)) & ~ (e22 = h5(e13)) & ~ (e22 = h5(e12)) & ~ (e22 = h5(e11)) & ~ (e22 = h5(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP278])])).
fof(f305, plain, ((~ (e23 = h5(e14)) & ~ (e23 = h5(e13)) & ~ (e23 = h5(e12)) & ~ (e23 = h5(e11)) & ~ (e23 = h5(e10))) | ~ sP279), inference(usedef, [], [e305])).
fof(e305, plain, (sP279 <=> (~ (e23 = h5(e14)) & ~ (e23 = h5(e13)) & ~ (e23 = h5(e12)) & ~ (e23 = h5(e11)) & ~ (e23 = h5(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP279])])).
fof(f23, plain, (((~ (e24 = h5(e14)) & ~ (e24 = h5(e13)) & ~ (e24 = h5(e12)) & ~ (e24 = h5(e11)) & ~ (e24 = h5(e10))) | (~ (e23 = h5(e14)) & ~ (e23 = h5(e13)) & ~ (e23 = h5(e12)) & ~ (e23 = h5(e11)) & ~ (e23 = h5(e10))) | (~ (e22 = h5(e14)) & ~ (e22 = h5(e13)) & ~ (e22 = h5(e12)) & ~ (e22 = h5(e11)) & ~ (e22 = h5(e10))) | (~ (e21 = h5(e14)) & ~ (e21 = h5(e13)) & ~ (e21 = h5(e12)) & ~ (e21 = h5(e11)) & ~ (e21 = h5(e10))) | (~ (e20 = h5(e14)) & ~ (e20 = h5(e13)) & ~ (e20 = h5(e12)) & ~ (e20 = h5(e11)) & ~ (e20 = h5(e10))) | ~ (h5(op1(e14, e14)) = op2(h5(e14), h5(e14))) | ~ (h5(op1(e14, e13)) = op2(h5(e14), h5(e13))) | ~ (h5(op1(e14, e12)) = op2(h5(e14), h5(e12))) | ~ (h5(op1(e14, e11)) = op2(h5(e14), h5(e11))) | ~ (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) | ~ (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))) & ((~ (e24 = h4(e14)) & ~ (e24 = h4(e13)) & ~ (e24 = h4(e12)) & ~ (e24 = h4(e11)) & ~ (e24 = h4(e10))) | (~ (e23 = h4(e14)) & ~ (e23 = h4(e13)) & ~ (e23 = h4(e12)) & ~ (e23 = h4(e11)) & ~ (e23 = h4(e10))) | (~ (e22 = h4(e14)) & ~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10))) | (~ (e21 = h4(e14)) & ~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10))) | (~ (e20 = h4(e14)) & ~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10))) | ~ (h4(op1(e14, e14)) = op2(h4(e14), h4(e14))) | ~ (h4(op1(e14, e13)) = op2(h4(e14), h4(e13))) | ~ (h4(op1(e14, e12)) = op2(h4(e14), h4(e12))) | ~ (h4(op1(e14, e11)) = op2(h4(e14), h4(e11))) | ~ (h4(op1(e14, e10)) = op2(h4(e14), h4(e10))) | ~ (h4(op1(e13, e14)) = op2(h4(e13), h4(e14))) | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e14)) = op2(h4(e12), h4(e14))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e14)) = op2(h4(e11), h4(e14))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e14)) = op2(h4(e10), h4(e14))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) & ((~ (e24 = h3(e14)) & ~ (e24 = h3(e13)) & ~ (e24 = h3(e12)) & ~ (e24 = h3(e11)) & ~ (e24 = h3(e10))) | (~ (e23 = h3(e14)) & ~ (e23 = h3(e13)) & ~ (e23 = h3(e12)) & ~ (e23 = h3(e11)) & ~ (e23 = h3(e10))) | (~ (e22 = h3(e14)) & ~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10))) | (~ (e21 = h3(e14)) & ~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10))) | (~ (e20 = h3(e14)) & ~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10))) | ~ (h3(op1(e14, e14)) = op2(h3(e14), h3(e14))) | ~ (h3(op1(e14, e13)) = op2(h3(e14), h3(e13))) | ~ (h3(op1(e14, e12)) = op2(h3(e14), h3(e12))) | ~ (h3(op1(e14, e11)) = op2(h3(e14), h3(e11))) | ~ (h3(op1(e14, e10)) = op2(h3(e14), h3(e10))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), h3(e14))) | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) & ((~ (e24 = h2(e14)) & ~ (e24 = h2(e13)) & ~ (e24 = h2(e12)) & ~ (e24 = h2(e11)) & ~ (e24 = h2(e10))) | (~ (e23 = h2(e14)) & ~ (e23 = h2(e13)) & ~ (e23 = h2(e12)) & ~ (e23 = h2(e11)) & ~ (e23 = h2(e10))) | (~ (e22 = h2(e14)) & ~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10))) | (~ (e21 = h2(e14)) & ~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10))) | (~ (e20 = h2(e14)) & ~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10))) | ~ (h2(op1(e14, e14)) = op2(h2(e14), h2(e14))) | ~ (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) | ~ (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) | ~ (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) | ~ (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) | ~ (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) & ((~ (e24 = h1(e14)) & ~ (e24 = h1(e13)) & ~ (e24 = h1(e12)) & ~ (e24 = h1(e11)) & ~ (e24 = h1(e10))) | (~ (e23 = h1(e14)) & ~ (e23 = h1(e13)) & ~ (e23 = h1(e12)) & ~ (e23 = h1(e11)) & ~ (e23 = h1(e10))) | (~ (e22 = h1(e14)) & ~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10))) | (~ (e21 = h1(e14)) & ~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10))) | (~ (e20 = h1(e14)) & ~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10))) | ~ (h1(op1(e14, e14)) = op2(h1(e14), h1(e14))) | ~ (h1(op1(e14, e13)) = op2(h1(e14), h1(e13))) | ~ (h1(op1(e14, e12)) = op2(h1(e14), h1(e12))) | ~ (h1(op1(e14, e11)) = op2(h1(e14), h1(e11))) | ~ (h1(op1(e14, e10)) = op2(h1(e14), h1(e10))) | ~ (h1(op1(e13, e14)) = op2(h1(e13), h1(e14))) | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e14)) = op2(h1(e12), h1(e14))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e14)) = op2(h1(e11), h1(e14))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e14)) = op2(h1(e10), h1(e14))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(ennf_transformation, [], [f22])).
fof(f22, plain, ~ ((((e24 = h5(e14)) | (e24 = h5(e13)) | (e24 = h5(e12)) | (e24 = h5(e11)) | (e24 = h5(e10))) & ((e23 = h5(e14)) | (e23 = h5(e13)) | (e23 = h5(e12)) | (e23 = h5(e11)) | (e23 = h5(e10))) & ((e22 = h5(e14)) | (e22 = h5(e13)) | (e22 = h5(e12)) | (e22 = h5(e11)) | (e22 = h5(e10))) & ((e21 = h5(e14)) | (e21 = h5(e13)) | (e21 = h5(e12)) | (e21 = h5(e11)) | (e21 = h5(e10))) & ((e20 = h5(e14)) | (e20 = h5(e13)) | (e20 = h5(e12)) | (e20 = h5(e11)) | (e20 = h5(e10))) & (h5(op1(e14, e14)) = op2(h5(e14), h5(e14))) & (h5(op1(e14, e13)) = op2(h5(e14), h5(e13))) & (h5(op1(e14, e12)) = op2(h5(e14), h5(e12))) & (h5(op1(e14, e11)) = op2(h5(e14), h5(e11))) & (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) & (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) & (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) & (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) & (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) & (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) & (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) & (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) & (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) & (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) & (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) & (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) & (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) & (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) & (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) & (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) & (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) & (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) & (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) & (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) & (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))) | (((e24 = h4(e14)) | (e24 = h4(e13)) | (e24 = h4(e12)) | (e24 = h4(e11)) | (e24 = h4(e10))) & ((e23 = h4(e14)) | (e23 = h4(e13)) | (e23 = h4(e12)) | (e23 = h4(e11)) | (e23 = h4(e10))) & ((e22 = h4(e14)) | (e22 = h4(e13)) | (e22 = h4(e12)) | (e22 = h4(e11)) | (e22 = h4(e10))) & ((e21 = h4(e14)) | (e21 = h4(e13)) | (e21 = h4(e12)) | (e21 = h4(e11)) | (e21 = h4(e10))) & ((e20 = h4(e14)) | (e20 = h4(e13)) | (e20 = h4(e12)) | (e20 = h4(e11)) | (e20 = h4(e10))) & (h4(op1(e14, e14)) = op2(h4(e14), h4(e14))) & (h4(op1(e14, e13)) = op2(h4(e14), h4(e13))) & (h4(op1(e14, e12)) = op2(h4(e14), h4(e12))) & (h4(op1(e14, e11)) = op2(h4(e14), h4(e11))) & (h4(op1(e14, e10)) = op2(h4(e14), h4(e10))) & (h4(op1(e13, e14)) = op2(h4(e13), h4(e14))) & (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) & (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) & (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) & (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) & (h4(op1(e12, e14)) = op2(h4(e12), h4(e14))) & (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) & (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) & (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) & (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) & (h4(op1(e11, e14)) = op2(h4(e11), h4(e14))) & (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) & (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) & (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) & (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) & (h4(op1(e10, e14)) = op2(h4(e10), h4(e14))) & (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) & (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) & (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) & (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) | (((e24 = h3(e14)) | (e24 = h3(e13)) | (e24 = h3(e12)) | (e24 = h3(e11)) | (e24 = h3(e10))) & ((e23 = h3(e14)) | (e23 = h3(e13)) | (e23 = h3(e12)) | (e23 = h3(e11)) | (e23 = h3(e10))) & ((e22 = h3(e14)) | (e22 = h3(e13)) | (e22 = h3(e12)) | (e22 = h3(e11)) | (e22 = h3(e10))) & ((e21 = h3(e14)) | (e21 = h3(e13)) | (e21 = h3(e12)) | (e21 = h3(e11)) | (e21 = h3(e10))) & ((e20 = h3(e14)) | (e20 = h3(e13)) | (e20 = h3(e12)) | (e20 = h3(e11)) | (e20 = h3(e10))) & (h3(op1(e14, e14)) = op2(h3(e14), h3(e14))) & (h3(op1(e14, e13)) = op2(h3(e14), h3(e13))) & (h3(op1(e14, e12)) = op2(h3(e14), h3(e12))) & (h3(op1(e14, e11)) = op2(h3(e14), h3(e11))) & (h3(op1(e14, e10)) = op2(h3(e14), h3(e10))) & (h3(op1(e13, e14)) = op2(h3(e13), h3(e14))) & (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) & (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) & (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) & (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) & (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) & (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) & (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) & (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) & (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) & (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) & (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) & (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) & (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) & (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) & (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) & (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) & (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) & (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) & (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) | (((e24 = h2(e14)) | (e24 = h2(e13)) | (e24 = h2(e12)) | (e24 = h2(e11)) | (e24 = h2(e10))) & ((e23 = h2(e14)) | (e23 = h2(e13)) | (e23 = h2(e12)) | (e23 = h2(e11)) | (e23 = h2(e10))) & ((e22 = h2(e14)) | (e22 = h2(e13)) | (e22 = h2(e12)) | (e22 = h2(e11)) | (e22 = h2(e10))) & ((e21 = h2(e14)) | (e21 = h2(e13)) | (e21 = h2(e12)) | (e21 = h2(e11)) | (e21 = h2(e10))) & ((e20 = h2(e14)) | (e20 = h2(e13)) | (e20 = h2(e12)) | (e20 = h2(e11)) | (e20 = h2(e10))) & (h2(op1(e14, e14)) = op2(h2(e14), h2(e14))) & (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) & (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) & (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) & (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) & (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) & (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) & (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) & (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) & (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) & (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) & (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) & (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) & (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) & (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) & (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) & (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) & (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) & (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) & (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) & (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) & (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) & (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) & (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) & (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) | (((e24 = h1(e14)) | (e24 = h1(e13)) | (e24 = h1(e12)) | (e24 = h1(e11)) | (e24 = h1(e10))) & ((e23 = h1(e14)) | (e23 = h1(e13)) | (e23 = h1(e12)) | (e23 = h1(e11)) | (e23 = h1(e10))) & ((e22 = h1(e14)) | (e22 = h1(e13)) | (e22 = h1(e12)) | (e22 = h1(e11)) | (e22 = h1(e10))) & ((e21 = h1(e14)) | (e21 = h1(e13)) | (e21 = h1(e12)) | (e21 = h1(e11)) | (e21 = h1(e10))) & ((e20 = h1(e14)) | (e20 = h1(e13)) | (e20 = h1(e12)) | (e20 = h1(e11)) | (e20 = h1(e10))) & (h1(op1(e14, e14)) = op2(h1(e14), h1(e14))) & (h1(op1(e14, e13)) = op2(h1(e14), h1(e13))) & (h1(op1(e14, e12)) = op2(h1(e14), h1(e12))) & (h1(op1(e14, e11)) = op2(h1(e14), h1(e11))) & (h1(op1(e14, e10)) = op2(h1(e14), h1(e10))) & (h1(op1(e13, e14)) = op2(h1(e13), h1(e14))) & (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) & (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) & (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) & (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) & (h1(op1(e12, e14)) = op2(h1(e12), h1(e14))) & (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) & (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) & (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) & (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) & (h1(op1(e11, e14)) = op2(h1(e11), h1(e14))) & (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) & (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) & (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) & (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) & (h1(op1(e10, e14)) = op2(h1(e10), h1(e14))) & (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) & (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) & (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) & (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(negated_conjecture, [], [f21])).
fof(f21, plain, ~ ((((e24 = h5(e14)) | (e24 = h5(e13)) | (e24 = h5(e12)) | (e24 = h5(e11)) | (e24 = h5(e10))) & ((e23 = h5(e14)) | (e23 = h5(e13)) | (e23 = h5(e12)) | (e23 = h5(e11)) | (e23 = h5(e10))) & ((e22 = h5(e14)) | (e22 = h5(e13)) | (e22 = h5(e12)) | (e22 = h5(e11)) | (e22 = h5(e10))) & ((e21 = h5(e14)) | (e21 = h5(e13)) | (e21 = h5(e12)) | (e21 = h5(e11)) | (e21 = h5(e10))) & ((e20 = h5(e14)) | (e20 = h5(e13)) | (e20 = h5(e12)) | (e20 = h5(e11)) | (e20 = h5(e10))) & (h5(op1(e14, e14)) = op2(h5(e14), h5(e14))) & (h5(op1(e14, e13)) = op2(h5(e14), h5(e13))) & (h5(op1(e14, e12)) = op2(h5(e14), h5(e12))) & (h5(op1(e14, e11)) = op2(h5(e14), h5(e11))) & (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) & (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) & (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) & (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) & (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) & (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) & (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) & (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) & (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) & (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) & (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) & (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) & (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) & (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) & (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) & (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) & (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) & (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) & (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) & (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) & (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))) | (((e24 = h4(e14)) | (e24 = h4(e13)) | (e24 = h4(e12)) | (e24 = h4(e11)) | (e24 = h4(e10))) & ((e23 = h4(e14)) | (e23 = h4(e13)) | (e23 = h4(e12)) | (e23 = h4(e11)) | (e23 = h4(e10))) & ((e22 = h4(e14)) | (e22 = h4(e13)) | (e22 = h4(e12)) | (e22 = h4(e11)) | (e22 = h4(e10))) & ((e21 = h4(e14)) | (e21 = h4(e13)) | (e21 = h4(e12)) | (e21 = h4(e11)) | (e21 = h4(e10))) & ((e20 = h4(e14)) | (e20 = h4(e13)) | (e20 = h4(e12)) | (e20 = h4(e11)) | (e20 = h4(e10))) & (h4(op1(e14, e14)) = op2(h4(e14), h4(e14))) & (h4(op1(e14, e13)) = op2(h4(e14), h4(e13))) & (h4(op1(e14, e12)) = op2(h4(e14), h4(e12))) & (h4(op1(e14, e11)) = op2(h4(e14), h4(e11))) & (h4(op1(e14, e10)) = op2(h4(e14), h4(e10))) & (h4(op1(e13, e14)) = op2(h4(e13), h4(e14))) & (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) & (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) & (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) & (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) & (h4(op1(e12, e14)) = op2(h4(e12), h4(e14))) & (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) & (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) & (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) & (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) & (h4(op1(e11, e14)) = op2(h4(e11), h4(e14))) & (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) & (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) & (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) & (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) & (h4(op1(e10, e14)) = op2(h4(e10), h4(e14))) & (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) & (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) & (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) & (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) | (((e24 = h3(e14)) | (e24 = h3(e13)) | (e24 = h3(e12)) | (e24 = h3(e11)) | (e24 = h3(e10))) & ((e23 = h3(e14)) | (e23 = h3(e13)) | (e23 = h3(e12)) | (e23 = h3(e11)) | (e23 = h3(e10))) & ((e22 = h3(e14)) | (e22 = h3(e13)) | (e22 = h3(e12)) | (e22 = h3(e11)) | (e22 = h3(e10))) & ((e21 = h3(e14)) | (e21 = h3(e13)) | (e21 = h3(e12)) | (e21 = h3(e11)) | (e21 = h3(e10))) & ((e20 = h3(e14)) | (e20 = h3(e13)) | (e20 = h3(e12)) | (e20 = h3(e11)) | (e20 = h3(e10))) & (h3(op1(e14, e14)) = op2(h3(e14), h3(e14))) & (h3(op1(e14, e13)) = op2(h3(e14), h3(e13))) & (h3(op1(e14, e12)) = op2(h3(e14), h3(e12))) & (h3(op1(e14, e11)) = op2(h3(e14), h3(e11))) & (h3(op1(e14, e10)) = op2(h3(e14), h3(e10))) & (h3(op1(e13, e14)) = op2(h3(e13), h3(e14))) & (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) & (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) & (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) & (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) & (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) & (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) & (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) & (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) & (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) & (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) & (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) & (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) & (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) & (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) & (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) & (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) & (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) & (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) & (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) | (((e24 = h2(e14)) | (e24 = h2(e13)) | (e24 = h2(e12)) | (e24 = h2(e11)) | (e24 = h2(e10))) & ((e23 = h2(e14)) | (e23 = h2(e13)) | (e23 = h2(e12)) | (e23 = h2(e11)) | (e23 = h2(e10))) & ((e22 = h2(e14)) | (e22 = h2(e13)) | (e22 = h2(e12)) | (e22 = h2(e11)) | (e22 = h2(e10))) & ((e21 = h2(e14)) | (e21 = h2(e13)) | (e21 = h2(e12)) | (e21 = h2(e11)) | (e21 = h2(e10))) & ((e20 = h2(e14)) | (e20 = h2(e13)) | (e20 = h2(e12)) | (e20 = h2(e11)) | (e20 = h2(e10))) & (h2(op1(e14, e14)) = op2(h2(e14), h2(e14))) & (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) & (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) & (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) & (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) & (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) & (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) & (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) & (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) & (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) & (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) & (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) & (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) & (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) & (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) & (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) & (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) & (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) & (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) & (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) & (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) & (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) & (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) & (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) & (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) | (((e24 = h1(e14)) | (e24 = h1(e13)) | (e24 = h1(e12)) | (e24 = h1(e11)) | (e24 = h1(e10))) & ((e23 = h1(e14)) | (e23 = h1(e13)) | (e23 = h1(e12)) | (e23 = h1(e11)) | (e23 = h1(e10))) & ((e22 = h1(e14)) | (e22 = h1(e13)) | (e22 = h1(e12)) | (e22 = h1(e11)) | (e22 = h1(e10))) & ((e21 = h1(e14)) | (e21 = h1(e13)) | (e21 = h1(e12)) | (e21 = h1(e11)) | (e21 = h1(e10))) & ((e20 = h1(e14)) | (e20 = h1(e13)) | (e20 = h1(e12)) | (e20 = h1(e11)) | (e20 = h1(e10))) & (h1(op1(e14, e14)) = op2(h1(e14), h1(e14))) & (h1(op1(e14, e13)) = op2(h1(e14), h1(e13))) & (h1(op1(e14, e12)) = op2(h1(e14), h1(e12))) & (h1(op1(e14, e11)) = op2(h1(e14), h1(e11))) & (h1(op1(e14, e10)) = op2(h1(e14), h1(e10))) & (h1(op1(e13, e14)) = op2(h1(e13), h1(e14))) & (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) & (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) & (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) & (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) & (h1(op1(e12, e14)) = op2(h1(e12), h1(e14))) & (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) & (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) & (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) & (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) & (h1(op1(e11, e14)) = op2(h1(e11), h1(e14))) & (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) & (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) & (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) & (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) & (h1(op1(e10, e14)) = op2(h1(e10), h1(e14))) & (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) & (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) & (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) & (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG051+1.p', co1)).
fof(f5192, plain, (~ spl280_570 | ~ spl280_572), inference(avatar_split_clause, [], [f5187, f5189, f5177])).
fof(f5187, plain, (~ (e20 = h2(e13)) | ~ sP268), inference(forward_demodulation, [], [f1844, f4893])).
fof(f1844, plain, (~ (e20 = h3(e10)) | ~ sP268), inference(cnf_transformation, [], [f578])).
fof(f578, plain, ((~ (e20 = h3(e14)) & ~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10))) | ~ sP268), inference(nnf_transformation, [], [f294])).
fof(f5168, plain, ~ spl280_567, inference(avatar_split_clause, [], [f5167, f5158])).
fof(f5167, plain, ~ sP269, inference(subsumption_resolution, [], [f1840, f4890])).
fof(f1840, plain, (~ (e21 = h3(e11)) | ~ sP269), inference(cnf_transformation, [], [f577])).
fof(f577, plain, ((~ (e21 = h3(e14)) & ~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10))) | ~ sP269), inference(nnf_transformation, [], [f295])).
fof(f5148, plain, ~ spl280_564, inference(avatar_split_clause, [], [f5147, f5139])).
fof(f5147, plain, ~ sP270, inference(subsumption_resolution, [], [f1836, f1774])).
fof(f1836, plain, (~ (e22 = h3(e12)) | ~ sP270), inference(cnf_transformation, [], [f576])).
fof(f576, plain, ((~ (e22 = h3(e14)) & ~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10))) | ~ sP270), inference(nnf_transformation, [], [f296])).
fof(f5128, plain, (~ spl280_561 | ~ spl280_562), inference(avatar_split_clause, [], [f1832, f5125, f5121])).
fof(f1832, plain, (~ (e23 = h3(e13)) | ~ sP271), inference(cnf_transformation, [], [f575])).
fof(f575, plain, ((~ (e23 = h3(e14)) & ~ (e23 = h3(e13)) & ~ (e23 = h3(e12)) & ~ (e23 = h3(e11)) & ~ (e23 = h3(e10))) | ~ sP271), inference(nnf_transformation, [], [f297])).
fof(f4853, plain, spl280_221, inference(avatar_split_clause, [], [f4852, f2889])).
fof(f4852, plain, (e20 = op2(e21, e21)), inference(forward_demodulation, [], [f1760, f1761])).
fof(f1760, plain, (e20 = op2(op2(op2(e22, e22), e22), op2(op2(e22, e22), e22))), inference(cnf_transformation, [], [f15])).
fof(f4851, plain, spl280_194, inference(avatar_split_clause, [], [f1762, f2775])).
fof(f1762, plain, (e23 = op2(e22, e22)), inference(cnf_transformation, [], [f15])).
fof(f4850, plain, spl280_91, inference(avatar_split_clause, [], [f4849, f2293])).
fof(f4849, plain, (e10 = op1(e11, e11)), inference(forward_demodulation, [], [f1756, f1757])).
fof(f1756, plain, (e10 = op1(op1(op1(e12, e12), e12), op1(op1(e12, e12), e12))), inference(cnf_transformation, [], [f14])).
fof(f4848, plain, spl280_64, inference(avatar_split_clause, [], [f1758, f2179])).
fof(f1758, plain, (e13 = op1(e12, e12)), inference(cnf_transformation, [], [f14])).
fof(f4847, plain, (spl280_495 | spl280_469 | spl280_443 | spl280_417 | spl280_391), inference(avatar_split_clause, [], [f1755, f3978, f4107, f4236, f4365, f4494])).
fof(f4494, plain, (spl280_495 <=> sP255), introduced(avatar_definition, [new_symbols(naming, [spl280_495])])).
fof(f4365, plain, (spl280_469 <=> sP256), introduced(avatar_definition, [new_symbols(naming, [spl280_469])])).
fof(f4236, plain, (spl280_443 <=> sP257), introduced(avatar_definition, [new_symbols(naming, [spl280_443])])).
fof(f4107, plain, (spl280_417 <=> sP258), introduced(avatar_definition, [new_symbols(naming, [spl280_417])])).
fof(f3978, plain, (spl280_391 <=> sP259), introduced(avatar_definition, [new_symbols(naming, [spl280_391])])).
fof(f1755, plain, (sP259 | sP258 | sP257 | sP256 | sP255), inference(cnf_transformation, [], [f285])).
fof(f285, plain, (sP259 | sP258 | sP257 | sP256 | sP255), inference(definition_folding, [], [f13, e284, e283, e282, e281, e280, e279, e278, e277, e276, e275, e274, e273, e272, e271, e270, e269, e268, e267, e266, e265, e264, e263, e262, e261, e260, e259, e258, e257, e256, e255, e254, e253, e252, e251, e250, e249, e248, e247, e246, e245, e244, e243, e242, e241, e240, e239, e238, e237, e236, e235, e234, e233, e232, e231, e230, e229, e228, e227, e226, e225, e224, e223, e222, e221, e220, e219, e218, e217, e216, e215, e214, e213, e212, e211, e210, e209, e208, e207, e206, e205, e204, e203, e202, e201, e200, e199, e198, e197, e196, e195, e194, e193, e192, e191, e190, e189, e188, e187, e186, e185, e184, e183, e182, e181, e180, e179, e178, e177, e176, e175, e174, e173, e172, e171, e170, e169, e168, e167, e166, e165, e164, e163, e162, e161, e160, e159, e158, e157, e156, e155])).
fof(f155, plain, ((~ (e20 = unit2) & (e20 = op2(e20, e20))) | ~ (e20 = op2(e20, e20)) | ~ sP130), inference(usedef, [], [e155])).
fof(e155, plain, (sP130 <=> ((~ (e20 = unit2) & (e20 = op2(e20, e20))) | ~ (e20 = op2(e20, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP130])])).
fof(f156, plain, ((~ (e20 = unit2) & (op2(e20, e20) = e21)) | ~ (e20 = op2(e20, e21)) | ~ sP131), inference(usedef, [], [e156])).
fof(e156, plain, (sP131 <=> ((~ (e20 = unit2) & (op2(e20, e20) = e21)) | ~ (e20 = op2(e20, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP131])])).
fof(f157, plain, ((~ (e20 = unit2) & (op2(e20, e20) = e22)) | ~ (e20 = op2(e20, e22)) | ~ sP132), inference(usedef, [], [e157])).
fof(e157, plain, (sP132 <=> ((~ (e20 = unit2) & (op2(e20, e20) = e22)) | ~ (e20 = op2(e20, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP132])])).
fof(f158, plain, ((~ (e20 = unit2) & (op2(e20, e20) = e23)) | ~ (e20 = op2(e20, e23)) | ~ sP133), inference(usedef, [], [e158])).
fof(e158, plain, (sP133 <=> ((~ (e20 = unit2) & (op2(e20, e20) = e23)) | ~ (e20 = op2(e20, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP133])])).
fof(f159, plain, ((~ (e20 = unit2) & (op2(e20, e20) = e24)) | ~ (e20 = op2(e20, e24)) | ~ sP134), inference(usedef, [], [e159])).
fof(e159, plain, (sP134 <=> ((~ (e20 = unit2) & (op2(e20, e20) = e24)) | ~ (e20 = op2(e20, e24)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP134])])).
fof(f160, plain, ((~ (e20 = unit2) & (e20 = op2(e21, e20))) | ~ (e20 = op2(e21, e20)) | ~ sP135), inference(usedef, [], [e160])).
fof(e160, plain, (sP135 <=> ((~ (e20 = unit2) & (e20 = op2(e21, e20))) | ~ (e20 = op2(e21, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP135])])).
fof(f161, plain, ((~ (e20 = unit2) & (e21 = op2(e21, e20))) | ~ (e20 = op2(e21, e21)) | ~ sP136), inference(usedef, [], [e161])).
fof(e161, plain, (sP136 <=> ((~ (e20 = unit2) & (e21 = op2(e21, e20))) | ~ (e20 = op2(e21, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP136])])).
fof(f162, plain, ((~ (e20 = unit2) & (e22 = op2(e21, e20))) | ~ (e20 = op2(e21, e22)) | ~ sP137), inference(usedef, [], [e162])).
fof(e162, plain, (sP137 <=> ((~ (e20 = unit2) & (e22 = op2(e21, e20))) | ~ (e20 = op2(e21, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP137])])).
fof(f163, plain, ((~ (e20 = unit2) & (e23 = op2(e21, e20))) | ~ (e20 = op2(e21, e23)) | ~ sP138), inference(usedef, [], [e163])).
fof(e163, plain, (sP138 <=> ((~ (e20 = unit2) & (e23 = op2(e21, e20))) | ~ (e20 = op2(e21, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP138])])).
fof(f164, plain, ((~ (e20 = unit2) & (e24 = op2(e21, e20))) | ~ (e20 = op2(e21, e24)) | ~ sP139), inference(usedef, [], [e164])).
fof(e164, plain, (sP139 <=> ((~ (e20 = unit2) & (e24 = op2(e21, e20))) | ~ (e20 = op2(e21, e24)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP139])])).
fof(f165, plain, ((~ (e20 = unit2) & (e20 = op2(e22, e20))) | ~ (e20 = op2(e22, e20)) | ~ sP140), inference(usedef, [], [e165])).
fof(e165, plain, (sP140 <=> ((~ (e20 = unit2) & (e20 = op2(e22, e20))) | ~ (e20 = op2(e22, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP140])])).
fof(f166, plain, ((~ (e20 = unit2) & (e21 = op2(e22, e20))) | ~ (e20 = op2(e22, e21)) | ~ sP141), inference(usedef, [], [e166])).
fof(e166, plain, (sP141 <=> ((~ (e20 = unit2) & (e21 = op2(e22, e20))) | ~ (e20 = op2(e22, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP141])])).
fof(f167, plain, ((~ (e20 = unit2) & (e22 = op2(e22, e20))) | ~ (e20 = op2(e22, e22)) | ~ sP142), inference(usedef, [], [e167])).
fof(e167, plain, (sP142 <=> ((~ (e20 = unit2) & (e22 = op2(e22, e20))) | ~ (e20 = op2(e22, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP142])])).
fof(f168, plain, ((~ (e20 = unit2) & (e23 = op2(e22, e20))) | ~ (e20 = op2(e22, e23)) | ~ sP143), inference(usedef, [], [e168])).
fof(e168, plain, (sP143 <=> ((~ (e20 = unit2) & (e23 = op2(e22, e20))) | ~ (e20 = op2(e22, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP143])])).
fof(f169, plain, ((~ (e20 = unit2) & (e24 = op2(e22, e20))) | ~ (e20 = op2(e22, e24)) | ~ sP144), inference(usedef, [], [e169])).
fof(e169, plain, (sP144 <=> ((~ (e20 = unit2) & (e24 = op2(e22, e20))) | ~ (e20 = op2(e22, e24)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP144])])).
fof(f170, plain, ((~ (e20 = unit2) & (e20 = op2(e23, e20))) | ~ (e20 = op2(e23, e20)) | ~ sP145), inference(usedef, [], [e170])).
fof(e170, plain, (sP145 <=> ((~ (e20 = unit2) & (e20 = op2(e23, e20))) | ~ (e20 = op2(e23, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP145])])).
fof(f171, plain, ((~ (e20 = unit2) & (e21 = op2(e23, e20))) | ~ (e20 = op2(e23, e21)) | ~ sP146), inference(usedef, [], [e171])).
fof(e171, plain, (sP146 <=> ((~ (e20 = unit2) & (e21 = op2(e23, e20))) | ~ (e20 = op2(e23, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP146])])).
fof(f172, plain, ((~ (e20 = unit2) & (e22 = op2(e23, e20))) | ~ (e20 = op2(e23, e22)) | ~ sP147), inference(usedef, [], [e172])).
fof(e172, plain, (sP147 <=> ((~ (e20 = unit2) & (e22 = op2(e23, e20))) | ~ (e20 = op2(e23, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP147])])).
fof(f173, plain, ((~ (e20 = unit2) & (e23 = op2(e23, e20))) | ~ (e20 = op2(e23, e23)) | ~ sP148), inference(usedef, [], [e173])).
fof(e173, plain, (sP148 <=> ((~ (e20 = unit2) & (e23 = op2(e23, e20))) | ~ (e20 = op2(e23, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP148])])).
fof(f174, plain, ((~ (e20 = unit2) & (e24 = op2(e23, e20))) | ~ (e20 = op2(e23, e24)) | ~ sP149), inference(usedef, [], [e174])).
fof(e174, plain, (sP149 <=> ((~ (e20 = unit2) & (e24 = op2(e23, e20))) | ~ (e20 = op2(e23, e24)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP149])])).
fof(f175, plain, ((~ (e20 = unit2) & (e20 = op2(e24, e20))) | ~ (e20 = op2(e24, e20)) | ~ sP150), inference(usedef, [], [e175])).
fof(e175, plain, (sP150 <=> ((~ (e20 = unit2) & (e20 = op2(e24, e20))) | ~ (e20 = op2(e24, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP150])])).
fof(f176, plain, ((~ (e20 = unit2) & (e21 = op2(e24, e20))) | ~ (e20 = op2(e24, e21)) | ~ sP151), inference(usedef, [], [e176])).
fof(e176, plain, (sP151 <=> ((~ (e20 = unit2) & (e21 = op2(e24, e20))) | ~ (e20 = op2(e24, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP151])])).
fof(f177, plain, ((~ (e20 = unit2) & (e22 = op2(e24, e20))) | ~ (e20 = op2(e24, e22)) | ~ sP152), inference(usedef, [], [e177])).
fof(e177, plain, (sP152 <=> ((~ (e20 = unit2) & (e22 = op2(e24, e20))) | ~ (e20 = op2(e24, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP152])])).
fof(f178, plain, ((~ (e20 = unit2) & (e23 = op2(e24, e20))) | ~ (e20 = op2(e24, e23)) | ~ sP153), inference(usedef, [], [e178])).
fof(e178, plain, (sP153 <=> ((~ (e20 = unit2) & (e23 = op2(e24, e20))) | ~ (e20 = op2(e24, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP153])])).
fof(f179, plain, ((~ (e20 = unit2) & (e24 = op2(e24, e20))) | ~ (e20 = op2(e24, e24)) | ~ sP154), inference(usedef, [], [e179])).
fof(e179, plain, (sP154 <=> ((~ (e20 = unit2) & (e24 = op2(e24, e20))) | ~ (e20 = op2(e24, e24)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP154])])).
fof(f180, plain, ((~ (e21 = unit2) & (e20 = op2(e20, e21))) | ~ (op2(e20, e20) = e21) | ~ sP155), inference(usedef, [], [e180])).
fof(e180, plain, (sP155 <=> ((~ (e21 = unit2) & (e20 = op2(e20, e21))) | ~ (op2(e20, e20) = e21))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP155])])).
fof(f181, plain, ((~ (e21 = unit2) & (e21 = op2(e20, e21))) | ~ (e21 = op2(e20, e21)) | ~ sP156), inference(usedef, [], [e181])).
fof(e181, plain, (sP156 <=> ((~ (e21 = unit2) & (e21 = op2(e20, e21))) | ~ (e21 = op2(e20, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP156])])).
fof(f182, plain, ((~ (e21 = unit2) & (e22 = op2(e20, e21))) | ~ (e21 = op2(e20, e22)) | ~ sP157), inference(usedef, [], [e182])).
fof(e182, plain, (sP157 <=> ((~ (e21 = unit2) & (e22 = op2(e20, e21))) | ~ (e21 = op2(e20, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP157])])).
fof(f183, plain, ((~ (e21 = unit2) & (e23 = op2(e20, e21))) | ~ (e21 = op2(e20, e23)) | ~ sP158), inference(usedef, [], [e183])).
fof(e183, plain, (sP158 <=> ((~ (e21 = unit2) & (e23 = op2(e20, e21))) | ~ (e21 = op2(e20, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP158])])).
fof(f184, plain, ((~ (e21 = unit2) & (e24 = op2(e20, e21))) | ~ (e21 = op2(e20, e24)) | ~ sP159), inference(usedef, [], [e184])).
fof(e184, plain, (sP159 <=> ((~ (e21 = unit2) & (e24 = op2(e20, e21))) | ~ (e21 = op2(e20, e24)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP159])])).
fof(f185, plain, ((~ (e21 = unit2) & (e20 = op2(e21, e21))) | ~ (e21 = op2(e21, e20)) | ~ sP160), inference(usedef, [], [e185])).
fof(e185, plain, (sP160 <=> ((~ (e21 = unit2) & (e20 = op2(e21, e21))) | ~ (e21 = op2(e21, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP160])])).
fof(f186, plain, ((~ (e21 = unit2) & (e21 = op2(e21, e21))) | ~ (e21 = op2(e21, e21)) | ~ sP161), inference(usedef, [], [e186])).
fof(e186, plain, (sP161 <=> ((~ (e21 = unit2) & (e21 = op2(e21, e21))) | ~ (e21 = op2(e21, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP161])])).
fof(f187, plain, ((~ (e21 = unit2) & (e22 = op2(e21, e21))) | ~ (e21 = op2(e21, e22)) | ~ sP162), inference(usedef, [], [e187])).
fof(e187, plain, (sP162 <=> ((~ (e21 = unit2) & (e22 = op2(e21, e21))) | ~ (e21 = op2(e21, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP162])])).
fof(f188, plain, ((~ (e21 = unit2) & (e23 = op2(e21, e21))) | ~ (e21 = op2(e21, e23)) | ~ sP163), inference(usedef, [], [e188])).
fof(e188, plain, (sP163 <=> ((~ (e21 = unit2) & (e23 = op2(e21, e21))) | ~ (e21 = op2(e21, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP163])])).
fof(f189, plain, ((~ (e21 = unit2) & (e24 = op2(e21, e21))) | ~ (e21 = op2(e21, e24)) | ~ sP164), inference(usedef, [], [e189])).
fof(e189, plain, (sP164 <=> ((~ (e21 = unit2) & (e24 = op2(e21, e21))) | ~ (e21 = op2(e21, e24)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP164])])).
fof(f190, plain, ((~ (e21 = unit2) & (e20 = op2(e22, e21))) | ~ (e21 = op2(e22, e20)) | ~ sP165), inference(usedef, [], [e190])).
fof(e190, plain, (sP165 <=> ((~ (e21 = unit2) & (e20 = op2(e22, e21))) | ~ (e21 = op2(e22, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP165])])).
fof(f191, plain, ((~ (e21 = unit2) & (e21 = op2(e22, e21))) | ~ (e21 = op2(e22, e21)) | ~ sP166), inference(usedef, [], [e191])).
fof(e191, plain, (sP166 <=> ((~ (e21 = unit2) & (e21 = op2(e22, e21))) | ~ (e21 = op2(e22, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP166])])).
fof(f192, plain, ((~ (e21 = unit2) & (e22 = op2(e22, e21))) | ~ (e21 = op2(e22, e22)) | ~ sP167), inference(usedef, [], [e192])).
fof(e192, plain, (sP167 <=> ((~ (e21 = unit2) & (e22 = op2(e22, e21))) | ~ (e21 = op2(e22, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP167])])).
fof(f193, plain, ((~ (e21 = unit2) & (e23 = op2(e22, e21))) | ~ (e21 = op2(e22, e23)) | ~ sP168), inference(usedef, [], [e193])).
fof(e193, plain, (sP168 <=> ((~ (e21 = unit2) & (e23 = op2(e22, e21))) | ~ (e21 = op2(e22, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP168])])).
fof(f194, plain, ((~ (e21 = unit2) & (e24 = op2(e22, e21))) | ~ (e21 = op2(e22, e24)) | ~ sP169), inference(usedef, [], [e194])).
fof(e194, plain, (sP169 <=> ((~ (e21 = unit2) & (e24 = op2(e22, e21))) | ~ (e21 = op2(e22, e24)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP169])])).
fof(f195, plain, ((~ (e21 = unit2) & (e20 = op2(e23, e21))) | ~ (e21 = op2(e23, e20)) | ~ sP170), inference(usedef, [], [e195])).
fof(e195, plain, (sP170 <=> ((~ (e21 = unit2) & (e20 = op2(e23, e21))) | ~ (e21 = op2(e23, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP170])])).
fof(f196, plain, ((~ (e21 = unit2) & (e21 = op2(e23, e21))) | ~ (e21 = op2(e23, e21)) | ~ sP171), inference(usedef, [], [e196])).
fof(e196, plain, (sP171 <=> ((~ (e21 = unit2) & (e21 = op2(e23, e21))) | ~ (e21 = op2(e23, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP171])])).
fof(f197, plain, ((~ (e21 = unit2) & (e22 = op2(e23, e21))) | ~ (e21 = op2(e23, e22)) | ~ sP172), inference(usedef, [], [e197])).
fof(e197, plain, (sP172 <=> ((~ (e21 = unit2) & (e22 = op2(e23, e21))) | ~ (e21 = op2(e23, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP172])])).
fof(f198, plain, ((~ (e21 = unit2) & (e23 = op2(e23, e21))) | ~ (e21 = op2(e23, e23)) | ~ sP173), inference(usedef, [], [e198])).
fof(e198, plain, (sP173 <=> ((~ (e21 = unit2) & (e23 = op2(e23, e21))) | ~ (e21 = op2(e23, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP173])])).
fof(f199, plain, ((~ (e21 = unit2) & (e24 = op2(e23, e21))) | ~ (e21 = op2(e23, e24)) | ~ sP174), inference(usedef, [], [e199])).
fof(e199, plain, (sP174 <=> ((~ (e21 = unit2) & (e24 = op2(e23, e21))) | ~ (e21 = op2(e23, e24)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP174])])).
fof(f200, plain, ((~ (e21 = unit2) & (e20 = op2(e24, e21))) | ~ (e21 = op2(e24, e20)) | ~ sP175), inference(usedef, [], [e200])).
fof(e200, plain, (sP175 <=> ((~ (e21 = unit2) & (e20 = op2(e24, e21))) | ~ (e21 = op2(e24, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP175])])).
fof(f201, plain, ((~ (e21 = unit2) & (e21 = op2(e24, e21))) | ~ (e21 = op2(e24, e21)) | ~ sP176), inference(usedef, [], [e201])).
fof(e201, plain, (sP176 <=> ((~ (e21 = unit2) & (e21 = op2(e24, e21))) | ~ (e21 = op2(e24, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP176])])).
fof(f202, plain, ((~ (e21 = unit2) & (e22 = op2(e24, e21))) | ~ (e21 = op2(e24, e22)) | ~ sP177), inference(usedef, [], [e202])).
fof(e202, plain, (sP177 <=> ((~ (e21 = unit2) & (e22 = op2(e24, e21))) | ~ (e21 = op2(e24, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP177])])).
fof(f203, plain, ((~ (e21 = unit2) & (e23 = op2(e24, e21))) | ~ (e21 = op2(e24, e23)) | ~ sP178), inference(usedef, [], [e203])).
fof(e203, plain, (sP178 <=> ((~ (e21 = unit2) & (e23 = op2(e24, e21))) | ~ (e21 = op2(e24, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP178])])).
fof(f204, plain, ((~ (e21 = unit2) & (e24 = op2(e24, e21))) | ~ (e21 = op2(e24, e24)) | ~ sP179), inference(usedef, [], [e204])).
fof(e204, plain, (sP179 <=> ((~ (e21 = unit2) & (e24 = op2(e24, e21))) | ~ (e21 = op2(e24, e24)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP179])])).
fof(f205, plain, ((~ (e22 = unit2) & (e20 = op2(e20, e22))) | ~ (op2(e20, e20) = e22) | ~ sP180), inference(usedef, [], [e205])).
fof(e205, plain, (sP180 <=> ((~ (e22 = unit2) & (e20 = op2(e20, e22))) | ~ (op2(e20, e20) = e22))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP180])])).
fof(f206, plain, ((~ (e22 = unit2) & (e21 = op2(e20, e22))) | ~ (e22 = op2(e20, e21)) | ~ sP181), inference(usedef, [], [e206])).
fof(e206, plain, (sP181 <=> ((~ (e22 = unit2) & (e21 = op2(e20, e22))) | ~ (e22 = op2(e20, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP181])])).
fof(f207, plain, ((~ (e22 = unit2) & (e22 = op2(e20, e22))) | ~ (e22 = op2(e20, e22)) | ~ sP182), inference(usedef, [], [e207])).
fof(e207, plain, (sP182 <=> ((~ (e22 = unit2) & (e22 = op2(e20, e22))) | ~ (e22 = op2(e20, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP182])])).
fof(f208, plain, ((~ (e22 = unit2) & (e23 = op2(e20, e22))) | ~ (e22 = op2(e20, e23)) | ~ sP183), inference(usedef, [], [e208])).
fof(e208, plain, (sP183 <=> ((~ (e22 = unit2) & (e23 = op2(e20, e22))) | ~ (e22 = op2(e20, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP183])])).
fof(f209, plain, ((~ (e22 = unit2) & (e24 = op2(e20, e22))) | ~ (e22 = op2(e20, e24)) | ~ sP184), inference(usedef, [], [e209])).
fof(e209, plain, (sP184 <=> ((~ (e22 = unit2) & (e24 = op2(e20, e22))) | ~ (e22 = op2(e20, e24)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP184])])).
fof(f210, plain, ((~ (e22 = unit2) & (e20 = op2(e21, e22))) | ~ (e22 = op2(e21, e20)) | ~ sP185), inference(usedef, [], [e210])).
fof(e210, plain, (sP185 <=> ((~ (e22 = unit2) & (e20 = op2(e21, e22))) | ~ (e22 = op2(e21, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP185])])).
fof(f211, plain, ((~ (e22 = unit2) & (e21 = op2(e21, e22))) | ~ (e22 = op2(e21, e21)) | ~ sP186), inference(usedef, [], [e211])).
fof(e211, plain, (sP186 <=> ((~ (e22 = unit2) & (e21 = op2(e21, e22))) | ~ (e22 = op2(e21, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP186])])).
fof(f212, plain, ((~ (e22 = unit2) & (e22 = op2(e21, e22))) | ~ (e22 = op2(e21, e22)) | ~ sP187), inference(usedef, [], [e212])).
fof(e212, plain, (sP187 <=> ((~ (e22 = unit2) & (e22 = op2(e21, e22))) | ~ (e22 = op2(e21, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP187])])).
fof(f213, plain, ((~ (e22 = unit2) & (e23 = op2(e21, e22))) | ~ (e22 = op2(e21, e23)) | ~ sP188), inference(usedef, [], [e213])).
fof(e213, plain, (sP188 <=> ((~ (e22 = unit2) & (e23 = op2(e21, e22))) | ~ (e22 = op2(e21, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP188])])).
fof(f214, plain, ((~ (e22 = unit2) & (e24 = op2(e21, e22))) | ~ (e22 = op2(e21, e24)) | ~ sP189), inference(usedef, [], [e214])).
fof(e214, plain, (sP189 <=> ((~ (e22 = unit2) & (e24 = op2(e21, e22))) | ~ (e22 = op2(e21, e24)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP189])])).
fof(f215, plain, ((~ (e22 = unit2) & (e20 = op2(e22, e22))) | ~ (e22 = op2(e22, e20)) | ~ sP190), inference(usedef, [], [e215])).
fof(e215, plain, (sP190 <=> ((~ (e22 = unit2) & (e20 = op2(e22, e22))) | ~ (e22 = op2(e22, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP190])])).
fof(f216, plain, ((~ (e22 = unit2) & (e21 = op2(e22, e22))) | ~ (e22 = op2(e22, e21)) | ~ sP191), inference(usedef, [], [e216])).
fof(e216, plain, (sP191 <=> ((~ (e22 = unit2) & (e21 = op2(e22, e22))) | ~ (e22 = op2(e22, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP191])])).
fof(f217, plain, ((~ (e22 = unit2) & (e22 = op2(e22, e22))) | ~ (e22 = op2(e22, e22)) | ~ sP192), inference(usedef, [], [e217])).
fof(e217, plain, (sP192 <=> ((~ (e22 = unit2) & (e22 = op2(e22, e22))) | ~ (e22 = op2(e22, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP192])])).
fof(f218, plain, ((~ (e22 = unit2) & (e23 = op2(e22, e22))) | ~ (e22 = op2(e22, e23)) | ~ sP193), inference(usedef, [], [e218])).
fof(e218, plain, (sP193 <=> ((~ (e22 = unit2) & (e23 = op2(e22, e22))) | ~ (e22 = op2(e22, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP193])])).
fof(f219, plain, ((~ (e22 = unit2) & (e24 = op2(e22, e22))) | ~ (e22 = op2(e22, e24)) | ~ sP194), inference(usedef, [], [e219])).
fof(e219, plain, (sP194 <=> ((~ (e22 = unit2) & (e24 = op2(e22, e22))) | ~ (e22 = op2(e22, e24)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP194])])).
fof(f220, plain, ((~ (e22 = unit2) & (e20 = op2(e23, e22))) | ~ (e22 = op2(e23, e20)) | ~ sP195), inference(usedef, [], [e220])).
fof(e220, plain, (sP195 <=> ((~ (e22 = unit2) & (e20 = op2(e23, e22))) | ~ (e22 = op2(e23, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP195])])).
fof(f221, plain, ((~ (e22 = unit2) & (e21 = op2(e23, e22))) | ~ (e22 = op2(e23, e21)) | ~ sP196), inference(usedef, [], [e221])).
fof(e221, plain, (sP196 <=> ((~ (e22 = unit2) & (e21 = op2(e23, e22))) | ~ (e22 = op2(e23, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP196])])).
fof(f222, plain, ((~ (e22 = unit2) & (e22 = op2(e23, e22))) | ~ (e22 = op2(e23, e22)) | ~ sP197), inference(usedef, [], [e222])).
fof(e222, plain, (sP197 <=> ((~ (e22 = unit2) & (e22 = op2(e23, e22))) | ~ (e22 = op2(e23, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP197])])).
fof(f223, plain, ((~ (e22 = unit2) & (e23 = op2(e23, e22))) | ~ (e22 = op2(e23, e23)) | ~ sP198), inference(usedef, [], [e223])).
fof(e223, plain, (sP198 <=> ((~ (e22 = unit2) & (e23 = op2(e23, e22))) | ~ (e22 = op2(e23, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP198])])).
fof(f224, plain, ((~ (e22 = unit2) & (e24 = op2(e23, e22))) | ~ (e22 = op2(e23, e24)) | ~ sP199), inference(usedef, [], [e224])).
fof(e224, plain, (sP199 <=> ((~ (e22 = unit2) & (e24 = op2(e23, e22))) | ~ (e22 = op2(e23, e24)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP199])])).
fof(f225, plain, ((~ (e22 = unit2) & (e20 = op2(e24, e22))) | ~ (e22 = op2(e24, e20)) | ~ sP200), inference(usedef, [], [e225])).
fof(e225, plain, (sP200 <=> ((~ (e22 = unit2) & (e20 = op2(e24, e22))) | ~ (e22 = op2(e24, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP200])])).
fof(f226, plain, ((~ (e22 = unit2) & (e21 = op2(e24, e22))) | ~ (e22 = op2(e24, e21)) | ~ sP201), inference(usedef, [], [e226])).
fof(e226, plain, (sP201 <=> ((~ (e22 = unit2) & (e21 = op2(e24, e22))) | ~ (e22 = op2(e24, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP201])])).
fof(f227, plain, ((~ (e22 = unit2) & (e22 = op2(e24, e22))) | ~ (e22 = op2(e24, e22)) | ~ sP202), inference(usedef, [], [e227])).
fof(e227, plain, (sP202 <=> ((~ (e22 = unit2) & (e22 = op2(e24, e22))) | ~ (e22 = op2(e24, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP202])])).
fof(f228, plain, ((~ (e22 = unit2) & (e23 = op2(e24, e22))) | ~ (e22 = op2(e24, e23)) | ~ sP203), inference(usedef, [], [e228])).
fof(e228, plain, (sP203 <=> ((~ (e22 = unit2) & (e23 = op2(e24, e22))) | ~ (e22 = op2(e24, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP203])])).
fof(f229, plain, ((~ (e22 = unit2) & (e24 = op2(e24, e22))) | ~ (e22 = op2(e24, e24)) | ~ sP204), inference(usedef, [], [e229])).
fof(e229, plain, (sP204 <=> ((~ (e22 = unit2) & (e24 = op2(e24, e22))) | ~ (e22 = op2(e24, e24)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP204])])).
fof(f230, plain, ((~ (e23 = unit2) & (e20 = op2(e20, e23))) | ~ (op2(e20, e20) = e23) | ~ sP205), inference(usedef, [], [e230])).
fof(e230, plain, (sP205 <=> ((~ (e23 = unit2) & (e20 = op2(e20, e23))) | ~ (op2(e20, e20) = e23))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP205])])).
fof(f231, plain, ((~ (e23 = unit2) & (e21 = op2(e20, e23))) | ~ (e23 = op2(e20, e21)) | ~ sP206), inference(usedef, [], [e231])).
fof(e231, plain, (sP206 <=> ((~ (e23 = unit2) & (e21 = op2(e20, e23))) | ~ (e23 = op2(e20, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP206])])).
fof(f232, plain, ((~ (e23 = unit2) & (e22 = op2(e20, e23))) | ~ (e23 = op2(e20, e22)) | ~ sP207), inference(usedef, [], [e232])).
fof(e232, plain, (sP207 <=> ((~ (e23 = unit2) & (e22 = op2(e20, e23))) | ~ (e23 = op2(e20, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP207])])).
fof(f233, plain, ((~ (e23 = unit2) & (e23 = op2(e20, e23))) | ~ (e23 = op2(e20, e23)) | ~ sP208), inference(usedef, [], [e233])).
fof(e233, plain, (sP208 <=> ((~ (e23 = unit2) & (e23 = op2(e20, e23))) | ~ (e23 = op2(e20, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP208])])).
fof(f234, plain, ((~ (e23 = unit2) & (e24 = op2(e20, e23))) | ~ (e23 = op2(e20, e24)) | ~ sP209), inference(usedef, [], [e234])).
fof(e234, plain, (sP209 <=> ((~ (e23 = unit2) & (e24 = op2(e20, e23))) | ~ (e23 = op2(e20, e24)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP209])])).
fof(f235, plain, ((~ (e23 = unit2) & (e20 = op2(e21, e23))) | ~ (e23 = op2(e21, e20)) | ~ sP210), inference(usedef, [], [e235])).
fof(e235, plain, (sP210 <=> ((~ (e23 = unit2) & (e20 = op2(e21, e23))) | ~ (e23 = op2(e21, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP210])])).
fof(f236, plain, ((~ (e23 = unit2) & (e21 = op2(e21, e23))) | ~ (e23 = op2(e21, e21)) | ~ sP211), inference(usedef, [], [e236])).
fof(e236, plain, (sP211 <=> ((~ (e23 = unit2) & (e21 = op2(e21, e23))) | ~ (e23 = op2(e21, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP211])])).
fof(f237, plain, ((~ (e23 = unit2) & (e22 = op2(e21, e23))) | ~ (e23 = op2(e21, e22)) | ~ sP212), inference(usedef, [], [e237])).
fof(e237, plain, (sP212 <=> ((~ (e23 = unit2) & (e22 = op2(e21, e23))) | ~ (e23 = op2(e21, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP212])])).
fof(f238, plain, ((~ (e23 = unit2) & (e23 = op2(e21, e23))) | ~ (e23 = op2(e21, e23)) | ~ sP213), inference(usedef, [], [e238])).
fof(e238, plain, (sP213 <=> ((~ (e23 = unit2) & (e23 = op2(e21, e23))) | ~ (e23 = op2(e21, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP213])])).
fof(f239, plain, ((~ (e23 = unit2) & (e24 = op2(e21, e23))) | ~ (e23 = op2(e21, e24)) | ~ sP214), inference(usedef, [], [e239])).
fof(e239, plain, (sP214 <=> ((~ (e23 = unit2) & (e24 = op2(e21, e23))) | ~ (e23 = op2(e21, e24)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP214])])).
fof(f240, plain, ((~ (e23 = unit2) & (e20 = op2(e22, e23))) | ~ (e23 = op2(e22, e20)) | ~ sP215), inference(usedef, [], [e240])).
fof(e240, plain, (sP215 <=> ((~ (e23 = unit2) & (e20 = op2(e22, e23))) | ~ (e23 = op2(e22, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP215])])).
fof(f241, plain, ((~ (e23 = unit2) & (e21 = op2(e22, e23))) | ~ (e23 = op2(e22, e21)) | ~ sP216), inference(usedef, [], [e241])).
fof(e241, plain, (sP216 <=> ((~ (e23 = unit2) & (e21 = op2(e22, e23))) | ~ (e23 = op2(e22, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP216])])).
fof(f242, plain, ((~ (e23 = unit2) & (e22 = op2(e22, e23))) | ~ (e23 = op2(e22, e22)) | ~ sP217), inference(usedef, [], [e242])).
fof(e242, plain, (sP217 <=> ((~ (e23 = unit2) & (e22 = op2(e22, e23))) | ~ (e23 = op2(e22, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP217])])).
fof(f243, plain, ((~ (e23 = unit2) & (e23 = op2(e22, e23))) | ~ (e23 = op2(e22, e23)) | ~ sP218), inference(usedef, [], [e243])).
fof(e243, plain, (sP218 <=> ((~ (e23 = unit2) & (e23 = op2(e22, e23))) | ~ (e23 = op2(e22, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP218])])).
fof(f244, plain, ((~ (e23 = unit2) & (e24 = op2(e22, e23))) | ~ (e23 = op2(e22, e24)) | ~ sP219), inference(usedef, [], [e244])).
fof(e244, plain, (sP219 <=> ((~ (e23 = unit2) & (e24 = op2(e22, e23))) | ~ (e23 = op2(e22, e24)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP219])])).
fof(f245, plain, ((~ (e23 = unit2) & (e20 = op2(e23, e23))) | ~ (e23 = op2(e23, e20)) | ~ sP220), inference(usedef, [], [e245])).
fof(e245, plain, (sP220 <=> ((~ (e23 = unit2) & (e20 = op2(e23, e23))) | ~ (e23 = op2(e23, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP220])])).
fof(f246, plain, ((~ (e23 = unit2) & (e21 = op2(e23, e23))) | ~ (e23 = op2(e23, e21)) | ~ sP221), inference(usedef, [], [e246])).
fof(e246, plain, (sP221 <=> ((~ (e23 = unit2) & (e21 = op2(e23, e23))) | ~ (e23 = op2(e23, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP221])])).
fof(f247, plain, ((~ (e23 = unit2) & (e22 = op2(e23, e23))) | ~ (e23 = op2(e23, e22)) | ~ sP222), inference(usedef, [], [e247])).
fof(e247, plain, (sP222 <=> ((~ (e23 = unit2) & (e22 = op2(e23, e23))) | ~ (e23 = op2(e23, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP222])])).
fof(f248, plain, ((~ (e23 = unit2) & (e23 = op2(e23, e23))) | ~ (e23 = op2(e23, e23)) | ~ sP223), inference(usedef, [], [e248])).
fof(e248, plain, (sP223 <=> ((~ (e23 = unit2) & (e23 = op2(e23, e23))) | ~ (e23 = op2(e23, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP223])])).
fof(f249, plain, ((~ (e23 = unit2) & (e24 = op2(e23, e23))) | ~ (e23 = op2(e23, e24)) | ~ sP224), inference(usedef, [], [e249])).
fof(e249, plain, (sP224 <=> ((~ (e23 = unit2) & (e24 = op2(e23, e23))) | ~ (e23 = op2(e23, e24)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP224])])).
fof(f250, plain, ((~ (e23 = unit2) & (e20 = op2(e24, e23))) | ~ (e23 = op2(e24, e20)) | ~ sP225), inference(usedef, [], [e250])).
fof(e250, plain, (sP225 <=> ((~ (e23 = unit2) & (e20 = op2(e24, e23))) | ~ (e23 = op2(e24, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP225])])).
fof(f251, plain, ((~ (e23 = unit2) & (e21 = op2(e24, e23))) | ~ (e23 = op2(e24, e21)) | ~ sP226), inference(usedef, [], [e251])).
fof(e251, plain, (sP226 <=> ((~ (e23 = unit2) & (e21 = op2(e24, e23))) | ~ (e23 = op2(e24, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP226])])).
fof(f252, plain, ((~ (e23 = unit2) & (e22 = op2(e24, e23))) | ~ (e23 = op2(e24, e22)) | ~ sP227), inference(usedef, [], [e252])).
fof(e252, plain, (sP227 <=> ((~ (e23 = unit2) & (e22 = op2(e24, e23))) | ~ (e23 = op2(e24, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP227])])).
fof(f253, plain, ((~ (e23 = unit2) & (e23 = op2(e24, e23))) | ~ (e23 = op2(e24, e23)) | ~ sP228), inference(usedef, [], [e253])).
fof(e253, plain, (sP228 <=> ((~ (e23 = unit2) & (e23 = op2(e24, e23))) | ~ (e23 = op2(e24, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP228])])).
fof(f254, plain, ((~ (e23 = unit2) & (e24 = op2(e24, e23))) | ~ (e23 = op2(e24, e24)) | ~ sP229), inference(usedef, [], [e254])).
fof(e254, plain, (sP229 <=> ((~ (e23 = unit2) & (e24 = op2(e24, e23))) | ~ (e23 = op2(e24, e24)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP229])])).
fof(f255, plain, ((~ (e24 = unit2) & (e20 = op2(e20, e24))) | ~ (op2(e20, e20) = e24) | ~ sP230), inference(usedef, [], [e255])).
fof(e255, plain, (sP230 <=> ((~ (e24 = unit2) & (e20 = op2(e20, e24))) | ~ (op2(e20, e20) = e24))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP230])])).
fof(f256, plain, ((~ (e24 = unit2) & (e21 = op2(e20, e24))) | ~ (e24 = op2(e20, e21)) | ~ sP231), inference(usedef, [], [e256])).
fof(e256, plain, (sP231 <=> ((~ (e24 = unit2) & (e21 = op2(e20, e24))) | ~ (e24 = op2(e20, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP231])])).
fof(f257, plain, ((~ (e24 = unit2) & (e22 = op2(e20, e24))) | ~ (e24 = op2(e20, e22)) | ~ sP232), inference(usedef, [], [e257])).
fof(e257, plain, (sP232 <=> ((~ (e24 = unit2) & (e22 = op2(e20, e24))) | ~ (e24 = op2(e20, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP232])])).
fof(f258, plain, ((~ (e24 = unit2) & (e23 = op2(e20, e24))) | ~ (e24 = op2(e20, e23)) | ~ sP233), inference(usedef, [], [e258])).
fof(e258, plain, (sP233 <=> ((~ (e24 = unit2) & (e23 = op2(e20, e24))) | ~ (e24 = op2(e20, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP233])])).
fof(f259, plain, ((~ (e24 = unit2) & (e24 = op2(e20, e24))) | ~ (e24 = op2(e20, e24)) | ~ sP234), inference(usedef, [], [e259])).
fof(e259, plain, (sP234 <=> ((~ (e24 = unit2) & (e24 = op2(e20, e24))) | ~ (e24 = op2(e20, e24)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP234])])).
fof(f260, plain, ((~ (e24 = unit2) & (e20 = op2(e21, e24))) | ~ (e24 = op2(e21, e20)) | ~ sP235), inference(usedef, [], [e260])).
fof(e260, plain, (sP235 <=> ((~ (e24 = unit2) & (e20 = op2(e21, e24))) | ~ (e24 = op2(e21, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP235])])).
fof(f261, plain, ((~ (e24 = unit2) & (e21 = op2(e21, e24))) | ~ (e24 = op2(e21, e21)) | ~ sP236), inference(usedef, [], [e261])).
fof(e261, plain, (sP236 <=> ((~ (e24 = unit2) & (e21 = op2(e21, e24))) | ~ (e24 = op2(e21, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP236])])).
fof(f262, plain, ((~ (e24 = unit2) & (e22 = op2(e21, e24))) | ~ (e24 = op2(e21, e22)) | ~ sP237), inference(usedef, [], [e262])).
fof(e262, plain, (sP237 <=> ((~ (e24 = unit2) & (e22 = op2(e21, e24))) | ~ (e24 = op2(e21, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP237])])).
fof(f263, plain, ((~ (e24 = unit2) & (e23 = op2(e21, e24))) | ~ (e24 = op2(e21, e23)) | ~ sP238), inference(usedef, [], [e263])).
fof(e263, plain, (sP238 <=> ((~ (e24 = unit2) & (e23 = op2(e21, e24))) | ~ (e24 = op2(e21, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP238])])).
fof(f264, plain, ((~ (e24 = unit2) & (e24 = op2(e21, e24))) | ~ (e24 = op2(e21, e24)) | ~ sP239), inference(usedef, [], [e264])).
fof(e264, plain, (sP239 <=> ((~ (e24 = unit2) & (e24 = op2(e21, e24))) | ~ (e24 = op2(e21, e24)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP239])])).
fof(f265, plain, ((~ (e24 = unit2) & (e20 = op2(e22, e24))) | ~ (e24 = op2(e22, e20)) | ~ sP240), inference(usedef, [], [e265])).
fof(e265, plain, (sP240 <=> ((~ (e24 = unit2) & (e20 = op2(e22, e24))) | ~ (e24 = op2(e22, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP240])])).
fof(f266, plain, ((~ (e24 = unit2) & (e21 = op2(e22, e24))) | ~ (e24 = op2(e22, e21)) | ~ sP241), inference(usedef, [], [e266])).
fof(e266, plain, (sP241 <=> ((~ (e24 = unit2) & (e21 = op2(e22, e24))) | ~ (e24 = op2(e22, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP241])])).
fof(f267, plain, ((~ (e24 = unit2) & (e22 = op2(e22, e24))) | ~ (e24 = op2(e22, e22)) | ~ sP242), inference(usedef, [], [e267])).
fof(e267, plain, (sP242 <=> ((~ (e24 = unit2) & (e22 = op2(e22, e24))) | ~ (e24 = op2(e22, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP242])])).
fof(f268, plain, ((~ (e24 = unit2) & (e23 = op2(e22, e24))) | ~ (e24 = op2(e22, e23)) | ~ sP243), inference(usedef, [], [e268])).
fof(e268, plain, (sP243 <=> ((~ (e24 = unit2) & (e23 = op2(e22, e24))) | ~ (e24 = op2(e22, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP243])])).
fof(f269, plain, ((~ (e24 = unit2) & (e24 = op2(e22, e24))) | ~ (e24 = op2(e22, e24)) | ~ sP244), inference(usedef, [], [e269])).
fof(e269, plain, (sP244 <=> ((~ (e24 = unit2) & (e24 = op2(e22, e24))) | ~ (e24 = op2(e22, e24)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP244])])).
fof(f270, plain, ((~ (e24 = unit2) & (e20 = op2(e23, e24))) | ~ (e24 = op2(e23, e20)) | ~ sP245), inference(usedef, [], [e270])).
fof(e270, plain, (sP245 <=> ((~ (e24 = unit2) & (e20 = op2(e23, e24))) | ~ (e24 = op2(e23, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP245])])).
fof(f271, plain, ((~ (e24 = unit2) & (e21 = op2(e23, e24))) | ~ (e24 = op2(e23, e21)) | ~ sP246), inference(usedef, [], [e271])).
fof(e271, plain, (sP246 <=> ((~ (e24 = unit2) & (e21 = op2(e23, e24))) | ~ (e24 = op2(e23, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP246])])).
fof(f272, plain, ((~ (e24 = unit2) & (e22 = op2(e23, e24))) | ~ (e24 = op2(e23, e22)) | ~ sP247), inference(usedef, [], [e272])).
fof(e272, plain, (sP247 <=> ((~ (e24 = unit2) & (e22 = op2(e23, e24))) | ~ (e24 = op2(e23, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP247])])).
fof(f273, plain, ((~ (e24 = unit2) & (e23 = op2(e23, e24))) | ~ (e24 = op2(e23, e23)) | ~ sP248), inference(usedef, [], [e273])).
fof(e273, plain, (sP248 <=> ((~ (e24 = unit2) & (e23 = op2(e23, e24))) | ~ (e24 = op2(e23, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP248])])).
fof(f274, plain, ((~ (e24 = unit2) & (e24 = op2(e23, e24))) | ~ (e24 = op2(e23, e24)) | ~ sP249), inference(usedef, [], [e274])).
fof(e274, plain, (sP249 <=> ((~ (e24 = unit2) & (e24 = op2(e23, e24))) | ~ (e24 = op2(e23, e24)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP249])])).
fof(f275, plain, ((~ (e24 = unit2) & (e20 = op2(e24, e24))) | ~ (e24 = op2(e24, e20)) | ~ sP250), inference(usedef, [], [e275])).
fof(e275, plain, (sP250 <=> ((~ (e24 = unit2) & (e20 = op2(e24, e24))) | ~ (e24 = op2(e24, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP250])])).
fof(f276, plain, ((~ (e24 = unit2) & (e21 = op2(e24, e24))) | ~ (e24 = op2(e24, e21)) | ~ sP251), inference(usedef, [], [e276])).
fof(e276, plain, (sP251 <=> ((~ (e24 = unit2) & (e21 = op2(e24, e24))) | ~ (e24 = op2(e24, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP251])])).
fof(f277, plain, ((~ (e24 = unit2) & (e22 = op2(e24, e24))) | ~ (e24 = op2(e24, e22)) | ~ sP252), inference(usedef, [], [e277])).
fof(e277, plain, (sP252 <=> ((~ (e24 = unit2) & (e22 = op2(e24, e24))) | ~ (e24 = op2(e24, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP252])])).
fof(f278, plain, ((~ (e24 = unit2) & (e23 = op2(e24, e24))) | ~ (e24 = op2(e24, e23)) | ~ sP253), inference(usedef, [], [e278])).
fof(e278, plain, (sP253 <=> ((~ (e24 = unit2) & (e23 = op2(e24, e24))) | ~ (e24 = op2(e24, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP253])])).
fof(f279, plain, ((~ (e24 = unit2) & (e24 = op2(e24, e24))) | ~ (e24 = op2(e24, e24)) | ~ sP254), inference(usedef, [], [e279])).
fof(e279, plain, (sP254 <=> ((~ (e24 = unit2) & (e24 = op2(e24, e24))) | ~ (e24 = op2(e24, e24)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP254])])).
fof(f280, plain, ((sP154 & sP153 & sP152 & sP151 & sP150 & sP149 & sP148 & sP147 & sP146 & sP145 & sP144 & sP143 & sP142 & sP141 & sP140 & sP139 & sP138 & sP137 & sP136 & sP135 & sP134 & sP133 & sP132 & sP131 & sP130) | ~ sP255), inference(usedef, [], [e280])).
fof(e280, plain, (sP255 <=> (sP154 & sP153 & sP152 & sP151 & sP150 & sP149 & sP148 & sP147 & sP146 & sP145 & sP144 & sP143 & sP142 & sP141 & sP140 & sP139 & sP138 & sP137 & sP136 & sP135 & sP134 & sP133 & sP132 & sP131 & sP130)), introduced(predicate_definition_introduction, [new_symbols(naming, [sP255])])).
fof(f281, plain, ((sP179 & sP178 & sP177 & sP176 & sP175 & sP174 & sP173 & sP172 & sP171 & sP170 & sP169 & sP168 & sP167 & sP166 & sP165 & sP164 & sP163 & sP162 & sP161 & sP160 & sP159 & sP158 & sP157 & sP156 & sP155) | ~ sP256), inference(usedef, [], [e281])).
fof(e281, plain, (sP256 <=> (sP179 & sP178 & sP177 & sP176 & sP175 & sP174 & sP173 & sP172 & sP171 & sP170 & sP169 & sP168 & sP167 & sP166 & sP165 & sP164 & sP163 & sP162 & sP161 & sP160 & sP159 & sP158 & sP157 & sP156 & sP155)), introduced(predicate_definition_introduction, [new_symbols(naming, [sP256])])).
fof(f282, plain, ((sP204 & sP203 & sP202 & sP201 & sP200 & sP199 & sP198 & sP197 & sP196 & sP195 & sP194 & sP193 & sP192 & sP191 & sP190 & sP189 & sP188 & sP187 & sP186 & sP185 & sP184 & sP183 & sP182 & sP181 & sP180) | ~ sP257), inference(usedef, [], [e282])).
fof(e282, plain, (sP257 <=> (sP204 & sP203 & sP202 & sP201 & sP200 & sP199 & sP198 & sP197 & sP196 & sP195 & sP194 & sP193 & sP192 & sP191 & sP190 & sP189 & sP188 & sP187 & sP186 & sP185 & sP184 & sP183 & sP182 & sP181 & sP180)), introduced(predicate_definition_introduction, [new_symbols(naming, [sP257])])).
fof(f283, plain, ((sP229 & sP228 & sP227 & sP226 & sP225 & sP224 & sP223 & sP222 & sP221 & sP220 & sP219 & sP218 & sP217 & sP216 & sP215 & sP214 & sP213 & sP212 & sP211 & sP210 & sP209 & sP208 & sP207 & sP206 & sP205) | ~ sP258), inference(usedef, [], [e283])).
fof(e283, plain, (sP258 <=> (sP229 & sP228 & sP227 & sP226 & sP225 & sP224 & sP223 & sP222 & sP221 & sP220 & sP219 & sP218 & sP217 & sP216 & sP215 & sP214 & sP213 & sP212 & sP211 & sP210 & sP209 & sP208 & sP207 & sP206 & sP205)), introduced(predicate_definition_introduction, [new_symbols(naming, [sP258])])).
fof(f284, plain, ((sP254 & sP253 & sP252 & sP251 & sP250 & sP249 & sP248 & sP247 & sP246 & sP245 & sP244 & sP243 & sP242 & sP241 & sP240 & sP239 & sP238 & sP237 & sP236 & sP235 & sP234 & sP233 & sP232 & sP231 & sP230) | ~ sP259), inference(usedef, [], [e284])).
fof(e284, plain, (sP259 <=> (sP254 & sP253 & sP252 & sP251 & sP250 & sP249 & sP248 & sP247 & sP246 & sP245 & sP244 & sP243 & sP242 & sP241 & sP240 & sP239 & sP238 & sP237 & sP236 & sP235 & sP234 & sP233 & sP232 & sP231 & sP230)), introduced(predicate_definition_introduction, [new_symbols(naming, [sP259])])).
fof(f13, plain, ((((~ (e24 = unit2) & (e24 = op2(e24, e24))) | ~ (e24 = op2(e24, e24))) & ((~ (e24 = unit2) & (e23 = op2(e24, e24))) | ~ (e24 = op2(e24, e23))) & ((~ (e24 = unit2) & (e22 = op2(e24, e24))) | ~ (e24 = op2(e24, e22))) & ((~ (e24 = unit2) & (e21 = op2(e24, e24))) | ~ (e24 = op2(e24, e21))) & ((~ (e24 = unit2) & (e20 = op2(e24, e24))) | ~ (e24 = op2(e24, e20))) & ((~ (e24 = unit2) & (e24 = op2(e23, e24))) | ~ (e24 = op2(e23, e24))) & ((~ (e24 = unit2) & (e23 = op2(e23, e24))) | ~ (e24 = op2(e23, e23))) & ((~ (e24 = unit2) & (e22 = op2(e23, e24))) | ~ (e24 = op2(e23, e22))) & ((~ (e24 = unit2) & (e21 = op2(e23, e24))) | ~ (e24 = op2(e23, e21))) & ((~ (e24 = unit2) & (e20 = op2(e23, e24))) | ~ (e24 = op2(e23, e20))) & ((~ (e24 = unit2) & (e24 = op2(e22, e24))) | ~ (e24 = op2(e22, e24))) & ((~ (e24 = unit2) & (e23 = op2(e22, e24))) | ~ (e24 = op2(e22, e23))) & ((~ (e24 = unit2) & (e22 = op2(e22, e24))) | ~ (e24 = op2(e22, e22))) & ((~ (e24 = unit2) & (e21 = op2(e22, e24))) | ~ (e24 = op2(e22, e21))) & ((~ (e24 = unit2) & (e20 = op2(e22, e24))) | ~ (e24 = op2(e22, e20))) & ((~ (e24 = unit2) & (e24 = op2(e21, e24))) | ~ (e24 = op2(e21, e24))) & ((~ (e24 = unit2) & (e23 = op2(e21, e24))) | ~ (e24 = op2(e21, e23))) & ((~ (e24 = unit2) & (e22 = op2(e21, e24))) | ~ (e24 = op2(e21, e22))) & ((~ (e24 = unit2) & (e21 = op2(e21, e24))) | ~ (e24 = op2(e21, e21))) & ((~ (e24 = unit2) & (e20 = op2(e21, e24))) | ~ (e24 = op2(e21, e20))) & ((~ (e24 = unit2) & (e24 = op2(e20, e24))) | ~ (e24 = op2(e20, e24))) & ((~ (e24 = unit2) & (e23 = op2(e20, e24))) | ~ (e24 = op2(e20, e23))) & ((~ (e24 = unit2) & (e22 = op2(e20, e24))) | ~ (e24 = op2(e20, e22))) & ((~ (e24 = unit2) & (e21 = op2(e20, e24))) | ~ (e24 = op2(e20, e21))) & ((~ (e24 = unit2) & (e20 = op2(e20, e24))) | ~ (op2(e20, e20) = e24))) | (((~ (e23 = unit2) & (e24 = op2(e24, e23))) | ~ (e23 = op2(e24, e24))) & ((~ (e23 = unit2) & (e23 = op2(e24, e23))) | ~ (e23 = op2(e24, e23))) & ((~ (e23 = unit2) & (e22 = op2(e24, e23))) | ~ (e23 = op2(e24, e22))) & ((~ (e23 = unit2) & (e21 = op2(e24, e23))) | ~ (e23 = op2(e24, e21))) & ((~ (e23 = unit2) & (e20 = op2(e24, e23))) | ~ (e23 = op2(e24, e20))) & ((~ (e23 = unit2) & (e24 = op2(e23, e23))) | ~ (e23 = op2(e23, e24))) & ((~ (e23 = unit2) & (e23 = op2(e23, e23))) | ~ (e23 = op2(e23, e23))) & ((~ (e23 = unit2) & (e22 = op2(e23, e23))) | ~ (e23 = op2(e23, e22))) & ((~ (e23 = unit2) & (e21 = op2(e23, e23))) | ~ (e23 = op2(e23, e21))) & ((~ (e23 = unit2) & (e20 = op2(e23, e23))) | ~ (e23 = op2(e23, e20))) & ((~ (e23 = unit2) & (e24 = op2(e22, e23))) | ~ (e23 = op2(e22, e24))) & ((~ (e23 = unit2) & (e23 = op2(e22, e23))) | ~ (e23 = op2(e22, e23))) & ((~ (e23 = unit2) & (e22 = op2(e22, e23))) | ~ (e23 = op2(e22, e22))) & ((~ (e23 = unit2) & (e21 = op2(e22, e23))) | ~ (e23 = op2(e22, e21))) & ((~ (e23 = unit2) & (e20 = op2(e22, e23))) | ~ (e23 = op2(e22, e20))) & ((~ (e23 = unit2) & (e24 = op2(e21, e23))) | ~ (e23 = op2(e21, e24))) & ((~ (e23 = unit2) & (e23 = op2(e21, e23))) | ~ (e23 = op2(e21, e23))) & ((~ (e23 = unit2) & (e22 = op2(e21, e23))) | ~ (e23 = op2(e21, e22))) & ((~ (e23 = unit2) & (e21 = op2(e21, e23))) | ~ (e23 = op2(e21, e21))) & ((~ (e23 = unit2) & (e20 = op2(e21, e23))) | ~ (e23 = op2(e21, e20))) & ((~ (e23 = unit2) & (e24 = op2(e20, e23))) | ~ (e23 = op2(e20, e24))) & ((~ (e23 = unit2) & (e23 = op2(e20, e23))) | ~ (e23 = op2(e20, e23))) & ((~ (e23 = unit2) & (e22 = op2(e20, e23))) | ~ (e23 = op2(e20, e22))) & ((~ (e23 = unit2) & (e21 = op2(e20, e23))) | ~ (e23 = op2(e20, e21))) & ((~ (e23 = unit2) & (e20 = op2(e20, e23))) | ~ (op2(e20, e20) = e23))) | (((~ (e22 = unit2) & (e24 = op2(e24, e22))) | ~ (e22 = op2(e24, e24))) & ((~ (e22 = unit2) & (e23 = op2(e24, e22))) | ~ (e22 = op2(e24, e23))) & ((~ (e22 = unit2) & (e22 = op2(e24, e22))) | ~ (e22 = op2(e24, e22))) & ((~ (e22 = unit2) & (e21 = op2(e24, e22))) | ~ (e22 = op2(e24, e21))) & ((~ (e22 = unit2) & (e20 = op2(e24, e22))) | ~ (e22 = op2(e24, e20))) & ((~ (e22 = unit2) & (e24 = op2(e23, e22))) | ~ (e22 = op2(e23, e24))) & ((~ (e22 = unit2) & (e23 = op2(e23, e22))) | ~ (e22 = op2(e23, e23))) & ((~ (e22 = unit2) & (e22 = op2(e23, e22))) | ~ (e22 = op2(e23, e22))) & ((~ (e22 = unit2) & (e21 = op2(e23, e22))) | ~ (e22 = op2(e23, e21))) & ((~ (e22 = unit2) & (e20 = op2(e23, e22))) | ~ (e22 = op2(e23, e20))) & ((~ (e22 = unit2) & (e24 = op2(e22, e22))) | ~ (e22 = op2(e22, e24))) & ((~ (e22 = unit2) & (e23 = op2(e22, e22))) | ~ (e22 = op2(e22, e23))) & ((~ (e22 = unit2) & (e22 = op2(e22, e22))) | ~ (e22 = op2(e22, e22))) & ((~ (e22 = unit2) & (e21 = op2(e22, e22))) | ~ (e22 = op2(e22, e21))) & ((~ (e22 = unit2) & (e20 = op2(e22, e22))) | ~ (e22 = op2(e22, e20))) & ((~ (e22 = unit2) & (e24 = op2(e21, e22))) | ~ (e22 = op2(e21, e24))) & ((~ (e22 = unit2) & (e23 = op2(e21, e22))) | ~ (e22 = op2(e21, e23))) & ((~ (e22 = unit2) & (e22 = op2(e21, e22))) | ~ (e22 = op2(e21, e22))) & ((~ (e22 = unit2) & (e21 = op2(e21, e22))) | ~ (e22 = op2(e21, e21))) & ((~ (e22 = unit2) & (e20 = op2(e21, e22))) | ~ (e22 = op2(e21, e20))) & ((~ (e22 = unit2) & (e24 = op2(e20, e22))) | ~ (e22 = op2(e20, e24))) & ((~ (e22 = unit2) & (e23 = op2(e20, e22))) | ~ (e22 = op2(e20, e23))) & ((~ (e22 = unit2) & (e22 = op2(e20, e22))) | ~ (e22 = op2(e20, e22))) & ((~ (e22 = unit2) & (e21 = op2(e20, e22))) | ~ (e22 = op2(e20, e21))) & ((~ (e22 = unit2) & (e20 = op2(e20, e22))) | ~ (op2(e20, e20) = e22))) | (((~ (e21 = unit2) & (e24 = op2(e24, e21))) | ~ (e21 = op2(e24, e24))) & ((~ (e21 = unit2) & (e23 = op2(e24, e21))) | ~ (e21 = op2(e24, e23))) & ((~ (e21 = unit2) & (e22 = op2(e24, e21))) | ~ (e21 = op2(e24, e22))) & ((~ (e21 = unit2) & (e21 = op2(e24, e21))) | ~ (e21 = op2(e24, e21))) & ((~ (e21 = unit2) & (e20 = op2(e24, e21))) | ~ (e21 = op2(e24, e20))) & ((~ (e21 = unit2) & (e24 = op2(e23, e21))) | ~ (e21 = op2(e23, e24))) & ((~ (e21 = unit2) & (e23 = op2(e23, e21))) | ~ (e21 = op2(e23, e23))) & ((~ (e21 = unit2) & (e22 = op2(e23, e21))) | ~ (e21 = op2(e23, e22))) & ((~ (e21 = unit2) & (e21 = op2(e23, e21))) | ~ (e21 = op2(e23, e21))) & ((~ (e21 = unit2) & (e20 = op2(e23, e21))) | ~ (e21 = op2(e23, e20))) & ((~ (e21 = unit2) & (e24 = op2(e22, e21))) | ~ (e21 = op2(e22, e24))) & ((~ (e21 = unit2) & (e23 = op2(e22, e21))) | ~ (e21 = op2(e22, e23))) & ((~ (e21 = unit2) & (e22 = op2(e22, e21))) | ~ (e21 = op2(e22, e22))) & ((~ (e21 = unit2) & (e21 = op2(e22, e21))) | ~ (e21 = op2(e22, e21))) & ((~ (e21 = unit2) & (e20 = op2(e22, e21))) | ~ (e21 = op2(e22, e20))) & ((~ (e21 = unit2) & (e24 = op2(e21, e21))) | ~ (e21 = op2(e21, e24))) & ((~ (e21 = unit2) & (e23 = op2(e21, e21))) | ~ (e21 = op2(e21, e23))) & ((~ (e21 = unit2) & (e22 = op2(e21, e21))) | ~ (e21 = op2(e21, e22))) & ((~ (e21 = unit2) & (e21 = op2(e21, e21))) | ~ (e21 = op2(e21, e21))) & ((~ (e21 = unit2) & (e20 = op2(e21, e21))) | ~ (e21 = op2(e21, e20))) & ((~ (e21 = unit2) & (e24 = op2(e20, e21))) | ~ (e21 = op2(e20, e24))) & ((~ (e21 = unit2) & (e23 = op2(e20, e21))) | ~ (e21 = op2(e20, e23))) & ((~ (e21 = unit2) & (e22 = op2(e20, e21))) | ~ (e21 = op2(e20, e22))) & ((~ (e21 = unit2) & (e21 = op2(e20, e21))) | ~ (e21 = op2(e20, e21))) & ((~ (e21 = unit2) & (e20 = op2(e20, e21))) | ~ (op2(e20, e20) = e21))) | (((~ (e20 = unit2) & (e24 = op2(e24, e20))) | ~ (e20 = op2(e24, e24))) & ((~ (e20 = unit2) & (e23 = op2(e24, e20))) | ~ (e20 = op2(e24, e23))) & ((~ (e20 = unit2) & (e22 = op2(e24, e20))) | ~ (e20 = op2(e24, e22))) & ((~ (e20 = unit2) & (e21 = op2(e24, e20))) | ~ (e20 = op2(e24, e21))) & ((~ (e20 = unit2) & (e20 = op2(e24, e20))) | ~ (e20 = op2(e24, e20))) & ((~ (e20 = unit2) & (e24 = op2(e23, e20))) | ~ (e20 = op2(e23, e24))) & ((~ (e20 = unit2) & (e23 = op2(e23, e20))) | ~ (e20 = op2(e23, e23))) & ((~ (e20 = unit2) & (e22 = op2(e23, e20))) | ~ (e20 = op2(e23, e22))) & ((~ (e20 = unit2) & (e21 = op2(e23, e20))) | ~ (e20 = op2(e23, e21))) & ((~ (e20 = unit2) & (e20 = op2(e23, e20))) | ~ (e20 = op2(e23, e20))) & ((~ (e20 = unit2) & (e24 = op2(e22, e20))) | ~ (e20 = op2(e22, e24))) & ((~ (e20 = unit2) & (e23 = op2(e22, e20))) | ~ (e20 = op2(e22, e23))) & ((~ (e20 = unit2) & (e22 = op2(e22, e20))) | ~ (e20 = op2(e22, e22))) & ((~ (e20 = unit2) & (e21 = op2(e22, e20))) | ~ (e20 = op2(e22, e21))) & ((~ (e20 = unit2) & (e20 = op2(e22, e20))) | ~ (e20 = op2(e22, e20))) & ((~ (e20 = unit2) & (e24 = op2(e21, e20))) | ~ (e20 = op2(e21, e24))) & ((~ (e20 = unit2) & (e23 = op2(e21, e20))) | ~ (e20 = op2(e21, e23))) & ((~ (e20 = unit2) & (e22 = op2(e21, e20))) | ~ (e20 = op2(e21, e22))) & ((~ (e20 = unit2) & (e21 = op2(e21, e20))) | ~ (e20 = op2(e21, e21))) & ((~ (e20 = unit2) & (e20 = op2(e21, e20))) | ~ (e20 = op2(e21, e20))) & ((~ (e20 = unit2) & (op2(e20, e20) = e24)) | ~ (e20 = op2(e20, e24))) & ((~ (e20 = unit2) & (op2(e20, e20) = e23)) | ~ (e20 = op2(e20, e23))) & ((~ (e20 = unit2) & (op2(e20, e20) = e22)) | ~ (e20 = op2(e20, e22))) & ((~ (e20 = unit2) & (op2(e20, e20) = e21)) | ~ (e20 = op2(e20, e21))) & ((~ (e20 = unit2) & (e20 = op2(e20, e20))) | ~ (e20 = op2(e20, e20))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG051+1.p', ax13)).
fof(f4835, plain, (~ spl280_514 | ~ spl280_221 | ~ spl280_256), inference(avatar_split_clause, [], [f1742, f3036, f2889, f4588])).
fof(f4588, plain, (spl280_514 <=> sP136), introduced(avatar_definition, [new_symbols(naming, [spl280_514])])).
fof(f1742, plain, (~ (e20 = unit2) | ~ (e20 = op2(e21, e21)) | ~ sP136), inference(cnf_transformation, [], [f560])).
fof(f560, plain, ((~ (e20 = unit2) & (e21 = op2(e21, e20))) | ~ (e20 = op2(e21, e21)) | ~ sP136), inference(nnf_transformation, [], [f161])).
fof(f4783, plain, (~ spl280_484 | ~ spl280_202 | spl280_196), inference(avatar_split_clause, [], [f1683, f2784, f2809, f4439])).
fof(f4439, plain, (spl280_484 <=> sP165), introduced(avatar_definition, [new_symbols(naming, [spl280_484])])).
fof(f2809, plain, (spl280_202 <=> (e21 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl280_202])])).
fof(f1683, plain, ((e20 = op2(e22, e21)) | ~ (e21 = op2(e22, e20)) | ~ sP165), inference(cnf_transformation, [], [f531])).
fof(f531, plain, ((~ (e21 = unit2) & (e20 = op2(e22, e21))) | ~ (e21 = op2(e22, e20)) | ~ sP165), inference(nnf_transformation, [], [f190])).
fof(f4778, plain, (~ spl280_481 | ~ spl280_187 | spl280_199), inference(avatar_split_clause, [], [f1677, f2796, f2746, f4424])).
fof(f4424, plain, (spl280_481 <=> sP168), introduced(avatar_definition, [new_symbols(naming, [spl280_481])])).
fof(f2746, plain, (spl280_187 <=> (e21 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl280_187])])).
fof(f1677, plain, ((e23 = op2(e22, e21)) | ~ (e21 = op2(e22, e23)) | ~ sP168), inference(cnf_transformation, [], [f528])).
fof(f528, plain, ((~ (e21 = unit2) & (e23 = op2(e22, e21))) | ~ (e21 = op2(e22, e23)) | ~ sP168), inference(nnf_transformation, [], [f193])).
fof(f4771, plain, (~ spl280_477 | ~ spl280_167 | spl280_173), inference(avatar_split_clause, [], [f1669, f2687, f2662, f4404])).
fof(f4404, plain, (spl280_477 <=> sP172), introduced(avatar_definition, [new_symbols(naming, [spl280_477])])).
fof(f1669, plain, ((e22 = op2(e23, e21)) | ~ (e21 = op2(e23, e22)) | ~ sP172), inference(cnf_transformation, [], [f524])).
fof(f524, plain, ((~ (e21 = unit2) & (e22 = op2(e23, e21))) | ~ (e21 = op2(e23, e22)) | ~ sP172), inference(nnf_transformation, [], [f197])).
fof(f4760, plain, (~ spl280_471 | ~ spl280_137 | spl280_149), inference(avatar_split_clause, [], [f1657, f2586, f2536, f4374])).
fof(f4374, plain, (spl280_471 <=> sP178), introduced(avatar_definition, [new_symbols(naming, [spl280_471])])).
fof(f1657, plain, ((e23 = op2(e24, e21)) | ~ (e21 = op2(e24, e23)) | ~ sP178), inference(cnf_transformation, [], [f518])).
fof(f518, plain, ((~ (e21 = unit2) & (e23 = op2(e24, e21))) | ~ (e21 = op2(e24, e23)) | ~ sP178), inference(nnf_transformation, [], [f203])).
fof(f4738, plain, (~ spl280_458 | ~ spl280_203 | spl280_191), inference(avatar_split_clause, [], [f1633, f2763, f2813, f4310])).
fof(f4310, plain, (spl280_458 <=> sP190), introduced(avatar_definition, [new_symbols(naming, [spl280_458])])).
fof(f1633, plain, ((e20 = op2(e22, e22)) | ~ (e22 = op2(e22, e20)) | ~ sP190), inference(cnf_transformation, [], [f506])).
fof(f506, plain, ((~ (e22 = unit2) & (e20 = op2(e22, e22))) | ~ (e22 = op2(e22, e20)) | ~ sP190), inference(nnf_transformation, [], [f215])).
fof(f4689, plain, (~ spl280_430 | ~ spl280_194 | spl280_188), inference(avatar_split_clause, [], [f1579, f2750, f2775, f4171])).
fof(f4171, plain, (spl280_430 <=> sP217), introduced(avatar_definition, [new_symbols(naming, [spl280_430])])).
fof(f1579, plain, ((e22 = op2(e22, e23)) | ~ (e23 = op2(e22, e22)) | ~ sP217), inference(cnf_transformation, [], [f479])).
fof(f479, plain, ((~ (e23 = unit2) & (e22 = op2(e22, e23))) | ~ (e23 = op2(e22, e22)) | ~ sP217), inference(nnf_transformation, [], [f242])).
fof(f4633, plain, (~ spl280_398 | ~ spl280_165 | spl280_159), inference(avatar_split_clause, [], [f1517, f2628, f2653, f4012])).
fof(f4012, plain, (spl280_398 <=> sP248), introduced(avatar_definition, [new_symbols(naming, [spl280_398])])).
fof(f1517, plain, ((e23 = op2(e23, e24)) | ~ (e24 = op2(e23, e23)) | ~ sP248), inference(cnf_transformation, [], [f448])).
fof(f448, plain, ((~ (e24 = unit2) & (e23 = op2(e23, e24))) | ~ (e24 = op2(e23, e23)) | ~ sP248), inference(nnf_transformation, [], [f273])).
fof(f4591, plain, (~ spl280_495 | spl280_514), inference(avatar_split_clause, [], [f1486, f4588, f4494])).
fof(f1486, plain, (sP136 | ~ sP255), inference(cnf_transformation, [], [f441])).
fof(f441, plain, ((sP154 & sP153 & sP152 & sP151 & sP150 & sP149 & sP148 & sP147 & sP146 & sP145 & sP144 & sP143 & sP142 & sP141 & sP140 & sP139 & sP138 & sP137 & sP136 & sP135 & sP134 & sP133 & sP132 & sP131 & sP130) | ~ sP255), inference(nnf_transformation, [], [f280])).
fof(f4442, plain, (~ spl280_469 | spl280_484), inference(avatar_split_clause, [], [f1465, f4439, f4365])).
fof(f1465, plain, (sP165 | ~ sP256), inference(cnf_transformation, [], [f440])).
fof(f440, plain, ((sP179 & sP178 & sP177 & sP176 & sP175 & sP174 & sP173 & sP172 & sP171 & sP170 & sP169 & sP168 & sP167 & sP166 & sP165 & sP164 & sP163 & sP162 & sP161 & sP160 & sP159 & sP158 & sP157 & sP156 & sP155) | ~ sP256), inference(nnf_transformation, [], [f281])).
fof(f4427, plain, (~ spl280_469 | spl280_481), inference(avatar_split_clause, [], [f1468, f4424, f4365])).
fof(f1468, plain, (sP168 | ~ sP256), inference(cnf_transformation, [], [f440])).
fof(f4407, plain, (~ spl280_469 | spl280_477), inference(avatar_split_clause, [], [f1472, f4404, f4365])).
fof(f1472, plain, (sP172 | ~ sP256), inference(cnf_transformation, [], [f440])).
fof(f4377, plain, (~ spl280_469 | spl280_471), inference(avatar_split_clause, [], [f1478, f4374, f4365])).
fof(f1478, plain, (sP178 | ~ sP256), inference(cnf_transformation, [], [f440])).
fof(f4313, plain, (~ spl280_443 | spl280_458), inference(avatar_split_clause, [], [f1440, f4310, f4236])).
fof(f1440, plain, (sP190 | ~ sP257), inference(cnf_transformation, [], [f439])).
fof(f439, plain, ((sP204 & sP203 & sP202 & sP201 & sP200 & sP199 & sP198 & sP197 & sP196 & sP195 & sP194 & sP193 & sP192 & sP191 & sP190 & sP189 & sP188 & sP187 & sP186 & sP185 & sP184 & sP183 & sP182 & sP181 & sP180) | ~ sP257), inference(nnf_transformation, [], [f282])).
fof(f4174, plain, (~ spl280_417 | spl280_430), inference(avatar_split_clause, [], [f1417, f4171, f4107])).
fof(f1417, plain, (sP217 | ~ sP258), inference(cnf_transformation, [], [f438])).
fof(f438, plain, ((sP229 & sP228 & sP227 & sP226 & sP225 & sP224 & sP223 & sP222 & sP221 & sP220 & sP219 & sP218 & sP217 & sP216 & sP215 & sP214 & sP213 & sP212 & sP211 & sP210 & sP209 & sP208 & sP207 & sP206 & sP205) | ~ sP258), inference(nnf_transformation, [], [f283])).
fof(f4015, plain, (~ spl280_391 | spl280_398), inference(avatar_split_clause, [], [f1398, f4012, f3978])).
fof(f1398, plain, (sP248 | ~ sP259), inference(cnf_transformation, [], [f437])).
fof(f437, plain, ((sP254 & sP253 & sP252 & sP251 & sP250 & sP249 & sP248 & sP247 & sP246 & sP245 & sP244 & sP243 & sP242 & sP241 & sP240 & sP239 & sP238 & sP237 & sP236 & sP235 & sP234 & sP233 & sP232 & sP231 & sP230) | ~ sP259), inference(nnf_transformation, [], [f284])).
fof(f3976, plain, (spl280_365 | spl280_339 | spl280_313 | spl280_287 | spl280_261), inference(avatar_split_clause, [], [f1379, f3107, f3236, f3365, f3494, f3623])).
fof(f3623, plain, (spl280_365 <=> sP125), introduced(avatar_definition, [new_symbols(naming, [spl280_365])])).
fof(f3494, plain, (spl280_339 <=> sP126), introduced(avatar_definition, [new_symbols(naming, [spl280_339])])).
fof(f3365, plain, (spl280_313 <=> sP127), introduced(avatar_definition, [new_symbols(naming, [spl280_313])])).
fof(f3236, plain, (spl280_287 <=> sP128), introduced(avatar_definition, [new_symbols(naming, [spl280_287])])).
fof(f3107, plain, (spl280_261 <=> sP129), introduced(avatar_definition, [new_symbols(naming, [spl280_261])])).
fof(f1379, plain, (sP129 | sP128 | sP127 | sP126 | sP125), inference(cnf_transformation, [], [f154])).
fof(f154, plain, (sP129 | sP128 | sP127 | sP126 | sP125), inference(definition_folding, [], [f12, e153, e152, e151, e150, e149, e148, e147, e146, e145, e144, e143, e142, e141, e140, e139, e138, e137, e136, e135, e134, e133, e132, e131, e130, e129, e128, e127, e126, e125, e124, e123, e122, e121, e120, e119, e118, e117, e116, e115, e114, e113, e112, e111, e110, e109, e108, e107, e106, e105, e104, e103, e102, e101, e100, e99, e98, e97, e96, e95, e94, e93, e92, e91, e90, e89, e88, e87, e86, e85, e84, e83, e82, e81, e80, e79, e78, e77, e76, e75, e74, e73, e72, e71, e70, e69, e68, e67, e66, e65, e64, e63, e62, e61, e60, e59, e58, e57, e56, e55, e54, e53, e52, e51, e50, e49, e48, e47, e46, e45, e44, e43, e42, e41, e40, e39, e38, e37, e36, e35, e34, e33, e32, e31, e30, e29, e28, e27, e26, e25, e24])).
fof(f24, plain, ((~ (e10 = unit1) & (e10 = op1(e10, e10))) | ~ (e10 = op1(e10, e10)) | ~ sP0), inference(usedef, [], [e24])).
fof(e24, plain, (sP0 <=> ((~ (e10 = unit1) & (e10 = op1(e10, e10))) | ~ (e10 = op1(e10, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f25, plain, ((~ (e10 = unit1) & (op1(e10, e10) = e11)) | ~ (e10 = op1(e10, e11)) | ~ sP1), inference(usedef, [], [e25])).
fof(e25, plain, (sP1 <=> ((~ (e10 = unit1) & (op1(e10, e10) = e11)) | ~ (e10 = op1(e10, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f26, plain, ((~ (e10 = unit1) & (op1(e10, e10) = e12)) | ~ (e10 = op1(e10, e12)) | ~ sP2), inference(usedef, [], [e26])).
fof(e26, plain, (sP2 <=> ((~ (e10 = unit1) & (op1(e10, e10) = e12)) | ~ (e10 = op1(e10, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f27, plain, ((~ (e10 = unit1) & (op1(e10, e10) = e13)) | ~ (e10 = op1(e10, e13)) | ~ sP3), inference(usedef, [], [e27])).
fof(e27, plain, (sP3 <=> ((~ (e10 = unit1) & (op1(e10, e10) = e13)) | ~ (e10 = op1(e10, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f28, plain, ((~ (e10 = unit1) & (op1(e10, e10) = e14)) | ~ (e10 = op1(e10, e14)) | ~ sP4), inference(usedef, [], [e28])).
fof(e28, plain, (sP4 <=> ((~ (e10 = unit1) & (op1(e10, e10) = e14)) | ~ (e10 = op1(e10, e14)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f29, plain, ((~ (e10 = unit1) & (e10 = op1(e11, e10))) | ~ (e10 = op1(e11, e10)) | ~ sP5), inference(usedef, [], [e29])).
fof(e29, plain, (sP5 <=> ((~ (e10 = unit1) & (e10 = op1(e11, e10))) | ~ (e10 = op1(e11, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f30, plain, ((~ (e10 = unit1) & (e11 = op1(e11, e10))) | ~ (e10 = op1(e11, e11)) | ~ sP6), inference(usedef, [], [e30])).
fof(e30, plain, (sP6 <=> ((~ (e10 = unit1) & (e11 = op1(e11, e10))) | ~ (e10 = op1(e11, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f31, plain, ((~ (e10 = unit1) & (e12 = op1(e11, e10))) | ~ (e10 = op1(e11, e12)) | ~ sP7), inference(usedef, [], [e31])).
fof(e31, plain, (sP7 <=> ((~ (e10 = unit1) & (e12 = op1(e11, e10))) | ~ (e10 = op1(e11, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f32, plain, ((~ (e10 = unit1) & (e13 = op1(e11, e10))) | ~ (e10 = op1(e11, e13)) | ~ sP8), inference(usedef, [], [e32])).
fof(e32, plain, (sP8 <=> ((~ (e10 = unit1) & (e13 = op1(e11, e10))) | ~ (e10 = op1(e11, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f33, plain, ((~ (e10 = unit1) & (e14 = op1(e11, e10))) | ~ (e10 = op1(e11, e14)) | ~ sP9), inference(usedef, [], [e33])).
fof(e33, plain, (sP9 <=> ((~ (e10 = unit1) & (e14 = op1(e11, e10))) | ~ (e10 = op1(e11, e14)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f34, plain, ((~ (e10 = unit1) & (e10 = op1(e12, e10))) | ~ (e10 = op1(e12, e10)) | ~ sP10), inference(usedef, [], [e34])).
fof(e34, plain, (sP10 <=> ((~ (e10 = unit1) & (e10 = op1(e12, e10))) | ~ (e10 = op1(e12, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f35, plain, ((~ (e10 = unit1) & (e11 = op1(e12, e10))) | ~ (e10 = op1(e12, e11)) | ~ sP11), inference(usedef, [], [e35])).
fof(e35, plain, (sP11 <=> ((~ (e10 = unit1) & (e11 = op1(e12, e10))) | ~ (e10 = op1(e12, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f36, plain, ((~ (e10 = unit1) & (e12 = op1(e12, e10))) | ~ (e10 = op1(e12, e12)) | ~ sP12), inference(usedef, [], [e36])).
fof(e36, plain, (sP12 <=> ((~ (e10 = unit1) & (e12 = op1(e12, e10))) | ~ (e10 = op1(e12, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f37, plain, ((~ (e10 = unit1) & (e13 = op1(e12, e10))) | ~ (e10 = op1(e12, e13)) | ~ sP13), inference(usedef, [], [e37])).
fof(e37, plain, (sP13 <=> ((~ (e10 = unit1) & (e13 = op1(e12, e10))) | ~ (e10 = op1(e12, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f38, plain, ((~ (e10 = unit1) & (e14 = op1(e12, e10))) | ~ (e10 = op1(e12, e14)) | ~ sP14), inference(usedef, [], [e38])).
fof(e38, plain, (sP14 <=> ((~ (e10 = unit1) & (e14 = op1(e12, e10))) | ~ (e10 = op1(e12, e14)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f39, plain, ((~ (e10 = unit1) & (e10 = op1(e13, e10))) | ~ (e10 = op1(e13, e10)) | ~ sP15), inference(usedef, [], [e39])).
fof(e39, plain, (sP15 <=> ((~ (e10 = unit1) & (e10 = op1(e13, e10))) | ~ (e10 = op1(e13, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f40, plain, ((~ (e10 = unit1) & (e11 = op1(e13, e10))) | ~ (e10 = op1(e13, e11)) | ~ sP16), inference(usedef, [], [e40])).
fof(e40, plain, (sP16 <=> ((~ (e10 = unit1) & (e11 = op1(e13, e10))) | ~ (e10 = op1(e13, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f41, plain, ((~ (e10 = unit1) & (e12 = op1(e13, e10))) | ~ (e10 = op1(e13, e12)) | ~ sP17), inference(usedef, [], [e41])).
fof(e41, plain, (sP17 <=> ((~ (e10 = unit1) & (e12 = op1(e13, e10))) | ~ (e10 = op1(e13, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f42, plain, ((~ (e10 = unit1) & (e13 = op1(e13, e10))) | ~ (e10 = op1(e13, e13)) | ~ sP18), inference(usedef, [], [e42])).
fof(e42, plain, (sP18 <=> ((~ (e10 = unit1) & (e13 = op1(e13, e10))) | ~ (e10 = op1(e13, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f43, plain, ((~ (e10 = unit1) & (e14 = op1(e13, e10))) | ~ (e10 = op1(e13, e14)) | ~ sP19), inference(usedef, [], [e43])).
fof(e43, plain, (sP19 <=> ((~ (e10 = unit1) & (e14 = op1(e13, e10))) | ~ (e10 = op1(e13, e14)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP19])])).
fof(f44, plain, ((~ (e10 = unit1) & (e10 = op1(e14, e10))) | ~ (e10 = op1(e14, e10)) | ~ sP20), inference(usedef, [], [e44])).
fof(e44, plain, (sP20 <=> ((~ (e10 = unit1) & (e10 = op1(e14, e10))) | ~ (e10 = op1(e14, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP20])])).
fof(f45, plain, ((~ (e10 = unit1) & (e11 = op1(e14, e10))) | ~ (e10 = op1(e14, e11)) | ~ sP21), inference(usedef, [], [e45])).
fof(e45, plain, (sP21 <=> ((~ (e10 = unit1) & (e11 = op1(e14, e10))) | ~ (e10 = op1(e14, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP21])])).
fof(f46, plain, ((~ (e10 = unit1) & (e12 = op1(e14, e10))) | ~ (e10 = op1(e14, e12)) | ~ sP22), inference(usedef, [], [e46])).
fof(e46, plain, (sP22 <=> ((~ (e10 = unit1) & (e12 = op1(e14, e10))) | ~ (e10 = op1(e14, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP22])])).
fof(f47, plain, ((~ (e10 = unit1) & (e13 = op1(e14, e10))) | ~ (e10 = op1(e14, e13)) | ~ sP23), inference(usedef, [], [e47])).
fof(e47, plain, (sP23 <=> ((~ (e10 = unit1) & (e13 = op1(e14, e10))) | ~ (e10 = op1(e14, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP23])])).
fof(f48, plain, ((~ (e10 = unit1) & (e14 = op1(e14, e10))) | ~ (e10 = op1(e14, e14)) | ~ sP24), inference(usedef, [], [e48])).
fof(e48, plain, (sP24 <=> ((~ (e10 = unit1) & (e14 = op1(e14, e10))) | ~ (e10 = op1(e14, e14)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP24])])).
fof(f49, plain, ((~ (e11 = unit1) & (e10 = op1(e10, e11))) | ~ (op1(e10, e10) = e11) | ~ sP25), inference(usedef, [], [e49])).
fof(e49, plain, (sP25 <=> ((~ (e11 = unit1) & (e10 = op1(e10, e11))) | ~ (op1(e10, e10) = e11))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP25])])).
fof(f50, plain, ((~ (e11 = unit1) & (e11 = op1(e10, e11))) | ~ (e11 = op1(e10, e11)) | ~ sP26), inference(usedef, [], [e50])).
fof(e50, plain, (sP26 <=> ((~ (e11 = unit1) & (e11 = op1(e10, e11))) | ~ (e11 = op1(e10, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP26])])).
fof(f51, plain, ((~ (e11 = unit1) & (e12 = op1(e10, e11))) | ~ (e11 = op1(e10, e12)) | ~ sP27), inference(usedef, [], [e51])).
fof(e51, plain, (sP27 <=> ((~ (e11 = unit1) & (e12 = op1(e10, e11))) | ~ (e11 = op1(e10, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP27])])).
fof(f52, plain, ((~ (e11 = unit1) & (e13 = op1(e10, e11))) | ~ (e11 = op1(e10, e13)) | ~ sP28), inference(usedef, [], [e52])).
fof(e52, plain, (sP28 <=> ((~ (e11 = unit1) & (e13 = op1(e10, e11))) | ~ (e11 = op1(e10, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP28])])).
fof(f53, plain, ((~ (e11 = unit1) & (e14 = op1(e10, e11))) | ~ (e11 = op1(e10, e14)) | ~ sP29), inference(usedef, [], [e53])).
fof(e53, plain, (sP29 <=> ((~ (e11 = unit1) & (e14 = op1(e10, e11))) | ~ (e11 = op1(e10, e14)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP29])])).
fof(f54, plain, ((~ (e11 = unit1) & (e10 = op1(e11, e11))) | ~ (e11 = op1(e11, e10)) | ~ sP30), inference(usedef, [], [e54])).
fof(e54, plain, (sP30 <=> ((~ (e11 = unit1) & (e10 = op1(e11, e11))) | ~ (e11 = op1(e11, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP30])])).
fof(f55, plain, ((~ (e11 = unit1) & (e11 = op1(e11, e11))) | ~ (e11 = op1(e11, e11)) | ~ sP31), inference(usedef, [], [e55])).
fof(e55, plain, (sP31 <=> ((~ (e11 = unit1) & (e11 = op1(e11, e11))) | ~ (e11 = op1(e11, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP31])])).
fof(f56, plain, ((~ (e11 = unit1) & (e12 = op1(e11, e11))) | ~ (e11 = op1(e11, e12)) | ~ sP32), inference(usedef, [], [e56])).
fof(e56, plain, (sP32 <=> ((~ (e11 = unit1) & (e12 = op1(e11, e11))) | ~ (e11 = op1(e11, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP32])])).
fof(f57, plain, ((~ (e11 = unit1) & (e13 = op1(e11, e11))) | ~ (e11 = op1(e11, e13)) | ~ sP33), inference(usedef, [], [e57])).
fof(e57, plain, (sP33 <=> ((~ (e11 = unit1) & (e13 = op1(e11, e11))) | ~ (e11 = op1(e11, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP33])])).
fof(f58, plain, ((~ (e11 = unit1) & (e14 = op1(e11, e11))) | ~ (e11 = op1(e11, e14)) | ~ sP34), inference(usedef, [], [e58])).
fof(e58, plain, (sP34 <=> ((~ (e11 = unit1) & (e14 = op1(e11, e11))) | ~ (e11 = op1(e11, e14)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP34])])).
fof(f59, plain, ((~ (e11 = unit1) & (e10 = op1(e12, e11))) | ~ (e11 = op1(e12, e10)) | ~ sP35), inference(usedef, [], [e59])).
fof(e59, plain, (sP35 <=> ((~ (e11 = unit1) & (e10 = op1(e12, e11))) | ~ (e11 = op1(e12, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP35])])).
fof(f60, plain, ((~ (e11 = unit1) & (e11 = op1(e12, e11))) | ~ (e11 = op1(e12, e11)) | ~ sP36), inference(usedef, [], [e60])).
fof(e60, plain, (sP36 <=> ((~ (e11 = unit1) & (e11 = op1(e12, e11))) | ~ (e11 = op1(e12, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP36])])).
fof(f61, plain, ((~ (e11 = unit1) & (e12 = op1(e12, e11))) | ~ (e11 = op1(e12, e12)) | ~ sP37), inference(usedef, [], [e61])).
fof(e61, plain, (sP37 <=> ((~ (e11 = unit1) & (e12 = op1(e12, e11))) | ~ (e11 = op1(e12, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP37])])).
fof(f62, plain, ((~ (e11 = unit1) & (e13 = op1(e12, e11))) | ~ (e11 = op1(e12, e13)) | ~ sP38), inference(usedef, [], [e62])).
fof(e62, plain, (sP38 <=> ((~ (e11 = unit1) & (e13 = op1(e12, e11))) | ~ (e11 = op1(e12, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP38])])).
fof(f63, plain, ((~ (e11 = unit1) & (e14 = op1(e12, e11))) | ~ (e11 = op1(e12, e14)) | ~ sP39), inference(usedef, [], [e63])).
fof(e63, plain, (sP39 <=> ((~ (e11 = unit1) & (e14 = op1(e12, e11))) | ~ (e11 = op1(e12, e14)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP39])])).
fof(f64, plain, ((~ (e11 = unit1) & (e10 = op1(e13, e11))) | ~ (e11 = op1(e13, e10)) | ~ sP40), inference(usedef, [], [e64])).
fof(e64, plain, (sP40 <=> ((~ (e11 = unit1) & (e10 = op1(e13, e11))) | ~ (e11 = op1(e13, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP40])])).
fof(f65, plain, ((~ (e11 = unit1) & (e11 = op1(e13, e11))) | ~ (e11 = op1(e13, e11)) | ~ sP41), inference(usedef, [], [e65])).
fof(e65, plain, (sP41 <=> ((~ (e11 = unit1) & (e11 = op1(e13, e11))) | ~ (e11 = op1(e13, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP41])])).
fof(f66, plain, ((~ (e11 = unit1) & (e12 = op1(e13, e11))) | ~ (e11 = op1(e13, e12)) | ~ sP42), inference(usedef, [], [e66])).
fof(e66, plain, (sP42 <=> ((~ (e11 = unit1) & (e12 = op1(e13, e11))) | ~ (e11 = op1(e13, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP42])])).
fof(f67, plain, ((~ (e11 = unit1) & (e13 = op1(e13, e11))) | ~ (e11 = op1(e13, e13)) | ~ sP43), inference(usedef, [], [e67])).
fof(e67, plain, (sP43 <=> ((~ (e11 = unit1) & (e13 = op1(e13, e11))) | ~ (e11 = op1(e13, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP43])])).
fof(f68, plain, ((~ (e11 = unit1) & (e14 = op1(e13, e11))) | ~ (e11 = op1(e13, e14)) | ~ sP44), inference(usedef, [], [e68])).
fof(e68, plain, (sP44 <=> ((~ (e11 = unit1) & (e14 = op1(e13, e11))) | ~ (e11 = op1(e13, e14)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP44])])).
fof(f69, plain, ((~ (e11 = unit1) & (e10 = op1(e14, e11))) | ~ (e11 = op1(e14, e10)) | ~ sP45), inference(usedef, [], [e69])).
fof(e69, plain, (sP45 <=> ((~ (e11 = unit1) & (e10 = op1(e14, e11))) | ~ (e11 = op1(e14, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP45])])).
fof(f70, plain, ((~ (e11 = unit1) & (e11 = op1(e14, e11))) | ~ (e11 = op1(e14, e11)) | ~ sP46), inference(usedef, [], [e70])).
fof(e70, plain, (sP46 <=> ((~ (e11 = unit1) & (e11 = op1(e14, e11))) | ~ (e11 = op1(e14, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP46])])).
fof(f71, plain, ((~ (e11 = unit1) & (e12 = op1(e14, e11))) | ~ (e11 = op1(e14, e12)) | ~ sP47), inference(usedef, [], [e71])).
fof(e71, plain, (sP47 <=> ((~ (e11 = unit1) & (e12 = op1(e14, e11))) | ~ (e11 = op1(e14, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP47])])).
fof(f72, plain, ((~ (e11 = unit1) & (e13 = op1(e14, e11))) | ~ (e11 = op1(e14, e13)) | ~ sP48), inference(usedef, [], [e72])).
fof(e72, plain, (sP48 <=> ((~ (e11 = unit1) & (e13 = op1(e14, e11))) | ~ (e11 = op1(e14, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP48])])).
fof(f73, plain, ((~ (e11 = unit1) & (e14 = op1(e14, e11))) | ~ (e11 = op1(e14, e14)) | ~ sP49), inference(usedef, [], [e73])).
fof(e73, plain, (sP49 <=> ((~ (e11 = unit1) & (e14 = op1(e14, e11))) | ~ (e11 = op1(e14, e14)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP49])])).
fof(f74, plain, ((~ (e12 = unit1) & (e10 = op1(e10, e12))) | ~ (op1(e10, e10) = e12) | ~ sP50), inference(usedef, [], [e74])).
fof(e74, plain, (sP50 <=> ((~ (e12 = unit1) & (e10 = op1(e10, e12))) | ~ (op1(e10, e10) = e12))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP50])])).
fof(f75, plain, ((~ (e12 = unit1) & (e11 = op1(e10, e12))) | ~ (e12 = op1(e10, e11)) | ~ sP51), inference(usedef, [], [e75])).
fof(e75, plain, (sP51 <=> ((~ (e12 = unit1) & (e11 = op1(e10, e12))) | ~ (e12 = op1(e10, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP51])])).
fof(f76, plain, ((~ (e12 = unit1) & (e12 = op1(e10, e12))) | ~ (e12 = op1(e10, e12)) | ~ sP52), inference(usedef, [], [e76])).
fof(e76, plain, (sP52 <=> ((~ (e12 = unit1) & (e12 = op1(e10, e12))) | ~ (e12 = op1(e10, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP52])])).
fof(f77, plain, ((~ (e12 = unit1) & (e13 = op1(e10, e12))) | ~ (e12 = op1(e10, e13)) | ~ sP53), inference(usedef, [], [e77])).
fof(e77, plain, (sP53 <=> ((~ (e12 = unit1) & (e13 = op1(e10, e12))) | ~ (e12 = op1(e10, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP53])])).
fof(f78, plain, ((~ (e12 = unit1) & (e14 = op1(e10, e12))) | ~ (e12 = op1(e10, e14)) | ~ sP54), inference(usedef, [], [e78])).
fof(e78, plain, (sP54 <=> ((~ (e12 = unit1) & (e14 = op1(e10, e12))) | ~ (e12 = op1(e10, e14)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP54])])).
fof(f79, plain, ((~ (e12 = unit1) & (e10 = op1(e11, e12))) | ~ (e12 = op1(e11, e10)) | ~ sP55), inference(usedef, [], [e79])).
fof(e79, plain, (sP55 <=> ((~ (e12 = unit1) & (e10 = op1(e11, e12))) | ~ (e12 = op1(e11, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP55])])).
fof(f80, plain, ((~ (e12 = unit1) & (e11 = op1(e11, e12))) | ~ (e12 = op1(e11, e11)) | ~ sP56), inference(usedef, [], [e80])).
fof(e80, plain, (sP56 <=> ((~ (e12 = unit1) & (e11 = op1(e11, e12))) | ~ (e12 = op1(e11, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP56])])).
fof(f81, plain, ((~ (e12 = unit1) & (e12 = op1(e11, e12))) | ~ (e12 = op1(e11, e12)) | ~ sP57), inference(usedef, [], [e81])).
fof(e81, plain, (sP57 <=> ((~ (e12 = unit1) & (e12 = op1(e11, e12))) | ~ (e12 = op1(e11, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP57])])).
fof(f82, plain, ((~ (e12 = unit1) & (e13 = op1(e11, e12))) | ~ (e12 = op1(e11, e13)) | ~ sP58), inference(usedef, [], [e82])).
fof(e82, plain, (sP58 <=> ((~ (e12 = unit1) & (e13 = op1(e11, e12))) | ~ (e12 = op1(e11, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP58])])).
fof(f83, plain, ((~ (e12 = unit1) & (e14 = op1(e11, e12))) | ~ (e12 = op1(e11, e14)) | ~ sP59), inference(usedef, [], [e83])).
fof(e83, plain, (sP59 <=> ((~ (e12 = unit1) & (e14 = op1(e11, e12))) | ~ (e12 = op1(e11, e14)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP59])])).
fof(f84, plain, ((~ (e12 = unit1) & (e10 = op1(e12, e12))) | ~ (e12 = op1(e12, e10)) | ~ sP60), inference(usedef, [], [e84])).
fof(e84, plain, (sP60 <=> ((~ (e12 = unit1) & (e10 = op1(e12, e12))) | ~ (e12 = op1(e12, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP60])])).
fof(f85, plain, ((~ (e12 = unit1) & (e11 = op1(e12, e12))) | ~ (e12 = op1(e12, e11)) | ~ sP61), inference(usedef, [], [e85])).
fof(e85, plain, (sP61 <=> ((~ (e12 = unit1) & (e11 = op1(e12, e12))) | ~ (e12 = op1(e12, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP61])])).
fof(f86, plain, ((~ (e12 = unit1) & (e12 = op1(e12, e12))) | ~ (e12 = op1(e12, e12)) | ~ sP62), inference(usedef, [], [e86])).
fof(e86, plain, (sP62 <=> ((~ (e12 = unit1) & (e12 = op1(e12, e12))) | ~ (e12 = op1(e12, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP62])])).
fof(f87, plain, ((~ (e12 = unit1) & (e13 = op1(e12, e12))) | ~ (e12 = op1(e12, e13)) | ~ sP63), inference(usedef, [], [e87])).
fof(e87, plain, (sP63 <=> ((~ (e12 = unit1) & (e13 = op1(e12, e12))) | ~ (e12 = op1(e12, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP63])])).
fof(f88, plain, ((~ (e12 = unit1) & (e14 = op1(e12, e12))) | ~ (e12 = op1(e12, e14)) | ~ sP64), inference(usedef, [], [e88])).
fof(e88, plain, (sP64 <=> ((~ (e12 = unit1) & (e14 = op1(e12, e12))) | ~ (e12 = op1(e12, e14)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP64])])).
fof(f89, plain, ((~ (e12 = unit1) & (e10 = op1(e13, e12))) | ~ (e12 = op1(e13, e10)) | ~ sP65), inference(usedef, [], [e89])).
fof(e89, plain, (sP65 <=> ((~ (e12 = unit1) & (e10 = op1(e13, e12))) | ~ (e12 = op1(e13, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP65])])).
fof(f90, plain, ((~ (e12 = unit1) & (e11 = op1(e13, e12))) | ~ (e12 = op1(e13, e11)) | ~ sP66), inference(usedef, [], [e90])).
fof(e90, plain, (sP66 <=> ((~ (e12 = unit1) & (e11 = op1(e13, e12))) | ~ (e12 = op1(e13, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP66])])).
fof(f91, plain, ((~ (e12 = unit1) & (e12 = op1(e13, e12))) | ~ (e12 = op1(e13, e12)) | ~ sP67), inference(usedef, [], [e91])).
fof(e91, plain, (sP67 <=> ((~ (e12 = unit1) & (e12 = op1(e13, e12))) | ~ (e12 = op1(e13, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP67])])).
fof(f92, plain, ((~ (e12 = unit1) & (e13 = op1(e13, e12))) | ~ (e12 = op1(e13, e13)) | ~ sP68), inference(usedef, [], [e92])).
fof(e92, plain, (sP68 <=> ((~ (e12 = unit1) & (e13 = op1(e13, e12))) | ~ (e12 = op1(e13, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP68])])).
fof(f93, plain, ((~ (e12 = unit1) & (e14 = op1(e13, e12))) | ~ (e12 = op1(e13, e14)) | ~ sP69), inference(usedef, [], [e93])).
fof(e93, plain, (sP69 <=> ((~ (e12 = unit1) & (e14 = op1(e13, e12))) | ~ (e12 = op1(e13, e14)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP69])])).
fof(f94, plain, ((~ (e12 = unit1) & (e10 = op1(e14, e12))) | ~ (e12 = op1(e14, e10)) | ~ sP70), inference(usedef, [], [e94])).
fof(e94, plain, (sP70 <=> ((~ (e12 = unit1) & (e10 = op1(e14, e12))) | ~ (e12 = op1(e14, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP70])])).
fof(f95, plain, ((~ (e12 = unit1) & (e11 = op1(e14, e12))) | ~ (e12 = op1(e14, e11)) | ~ sP71), inference(usedef, [], [e95])).
fof(e95, plain, (sP71 <=> ((~ (e12 = unit1) & (e11 = op1(e14, e12))) | ~ (e12 = op1(e14, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP71])])).
fof(f96, plain, ((~ (e12 = unit1) & (e12 = op1(e14, e12))) | ~ (e12 = op1(e14, e12)) | ~ sP72), inference(usedef, [], [e96])).
fof(e96, plain, (sP72 <=> ((~ (e12 = unit1) & (e12 = op1(e14, e12))) | ~ (e12 = op1(e14, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP72])])).
fof(f97, plain, ((~ (e12 = unit1) & (e13 = op1(e14, e12))) | ~ (e12 = op1(e14, e13)) | ~ sP73), inference(usedef, [], [e97])).
fof(e97, plain, (sP73 <=> ((~ (e12 = unit1) & (e13 = op1(e14, e12))) | ~ (e12 = op1(e14, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP73])])).
fof(f98, plain, ((~ (e12 = unit1) & (e14 = op1(e14, e12))) | ~ (e12 = op1(e14, e14)) | ~ sP74), inference(usedef, [], [e98])).
fof(e98, plain, (sP74 <=> ((~ (e12 = unit1) & (e14 = op1(e14, e12))) | ~ (e12 = op1(e14, e14)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP74])])).
fof(f99, plain, ((~ (e13 = unit1) & (e10 = op1(e10, e13))) | ~ (op1(e10, e10) = e13) | ~ sP75), inference(usedef, [], [e99])).
fof(e99, plain, (sP75 <=> ((~ (e13 = unit1) & (e10 = op1(e10, e13))) | ~ (op1(e10, e10) = e13))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP75])])).
fof(f100, plain, ((~ (e13 = unit1) & (e11 = op1(e10, e13))) | ~ (e13 = op1(e10, e11)) | ~ sP76), inference(usedef, [], [e100])).
fof(e100, plain, (sP76 <=> ((~ (e13 = unit1) & (e11 = op1(e10, e13))) | ~ (e13 = op1(e10, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP76])])).
fof(f101, plain, ((~ (e13 = unit1) & (e12 = op1(e10, e13))) | ~ (e13 = op1(e10, e12)) | ~ sP77), inference(usedef, [], [e101])).
fof(e101, plain, (sP77 <=> ((~ (e13 = unit1) & (e12 = op1(e10, e13))) | ~ (e13 = op1(e10, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP77])])).
fof(f102, plain, ((~ (e13 = unit1) & (e13 = op1(e10, e13))) | ~ (e13 = op1(e10, e13)) | ~ sP78), inference(usedef, [], [e102])).
fof(e102, plain, (sP78 <=> ((~ (e13 = unit1) & (e13 = op1(e10, e13))) | ~ (e13 = op1(e10, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP78])])).
fof(f103, plain, ((~ (e13 = unit1) & (e14 = op1(e10, e13))) | ~ (e13 = op1(e10, e14)) | ~ sP79), inference(usedef, [], [e103])).
fof(e103, plain, (sP79 <=> ((~ (e13 = unit1) & (e14 = op1(e10, e13))) | ~ (e13 = op1(e10, e14)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP79])])).
fof(f104, plain, ((~ (e13 = unit1) & (e10 = op1(e11, e13))) | ~ (e13 = op1(e11, e10)) | ~ sP80), inference(usedef, [], [e104])).
fof(e104, plain, (sP80 <=> ((~ (e13 = unit1) & (e10 = op1(e11, e13))) | ~ (e13 = op1(e11, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP80])])).
fof(f105, plain, ((~ (e13 = unit1) & (e11 = op1(e11, e13))) | ~ (e13 = op1(e11, e11)) | ~ sP81), inference(usedef, [], [e105])).
fof(e105, plain, (sP81 <=> ((~ (e13 = unit1) & (e11 = op1(e11, e13))) | ~ (e13 = op1(e11, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP81])])).
fof(f106, plain, ((~ (e13 = unit1) & (e12 = op1(e11, e13))) | ~ (e13 = op1(e11, e12)) | ~ sP82), inference(usedef, [], [e106])).
fof(e106, plain, (sP82 <=> ((~ (e13 = unit1) & (e12 = op1(e11, e13))) | ~ (e13 = op1(e11, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP82])])).
fof(f107, plain, ((~ (e13 = unit1) & (e13 = op1(e11, e13))) | ~ (e13 = op1(e11, e13)) | ~ sP83), inference(usedef, [], [e107])).
fof(e107, plain, (sP83 <=> ((~ (e13 = unit1) & (e13 = op1(e11, e13))) | ~ (e13 = op1(e11, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP83])])).
fof(f108, plain, ((~ (e13 = unit1) & (e14 = op1(e11, e13))) | ~ (e13 = op1(e11, e14)) | ~ sP84), inference(usedef, [], [e108])).
fof(e108, plain, (sP84 <=> ((~ (e13 = unit1) & (e14 = op1(e11, e13))) | ~ (e13 = op1(e11, e14)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP84])])).
fof(f109, plain, ((~ (e13 = unit1) & (e10 = op1(e12, e13))) | ~ (e13 = op1(e12, e10)) | ~ sP85), inference(usedef, [], [e109])).
fof(e109, plain, (sP85 <=> ((~ (e13 = unit1) & (e10 = op1(e12, e13))) | ~ (e13 = op1(e12, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP85])])).
fof(f110, plain, ((~ (e13 = unit1) & (e11 = op1(e12, e13))) | ~ (e13 = op1(e12, e11)) | ~ sP86), inference(usedef, [], [e110])).
fof(e110, plain, (sP86 <=> ((~ (e13 = unit1) & (e11 = op1(e12, e13))) | ~ (e13 = op1(e12, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP86])])).
fof(f111, plain, ((~ (e13 = unit1) & (e12 = op1(e12, e13))) | ~ (e13 = op1(e12, e12)) | ~ sP87), inference(usedef, [], [e111])).
fof(e111, plain, (sP87 <=> ((~ (e13 = unit1) & (e12 = op1(e12, e13))) | ~ (e13 = op1(e12, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP87])])).
fof(f112, plain, ((~ (e13 = unit1) & (e13 = op1(e12, e13))) | ~ (e13 = op1(e12, e13)) | ~ sP88), inference(usedef, [], [e112])).
fof(e112, plain, (sP88 <=> ((~ (e13 = unit1) & (e13 = op1(e12, e13))) | ~ (e13 = op1(e12, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP88])])).
fof(f113, plain, ((~ (e13 = unit1) & (e14 = op1(e12, e13))) | ~ (e13 = op1(e12, e14)) | ~ sP89), inference(usedef, [], [e113])).
fof(e113, plain, (sP89 <=> ((~ (e13 = unit1) & (e14 = op1(e12, e13))) | ~ (e13 = op1(e12, e14)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP89])])).
fof(f114, plain, ((~ (e13 = unit1) & (e10 = op1(e13, e13))) | ~ (e13 = op1(e13, e10)) | ~ sP90), inference(usedef, [], [e114])).
fof(e114, plain, (sP90 <=> ((~ (e13 = unit1) & (e10 = op1(e13, e13))) | ~ (e13 = op1(e13, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP90])])).
fof(f115, plain, ((~ (e13 = unit1) & (e11 = op1(e13, e13))) | ~ (e13 = op1(e13, e11)) | ~ sP91), inference(usedef, [], [e115])).
fof(e115, plain, (sP91 <=> ((~ (e13 = unit1) & (e11 = op1(e13, e13))) | ~ (e13 = op1(e13, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP91])])).
fof(f116, plain, ((~ (e13 = unit1) & (e12 = op1(e13, e13))) | ~ (e13 = op1(e13, e12)) | ~ sP92), inference(usedef, [], [e116])).
fof(e116, plain, (sP92 <=> ((~ (e13 = unit1) & (e12 = op1(e13, e13))) | ~ (e13 = op1(e13, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP92])])).
fof(f117, plain, ((~ (e13 = unit1) & (e13 = op1(e13, e13))) | ~ (e13 = op1(e13, e13)) | ~ sP93), inference(usedef, [], [e117])).
fof(e117, plain, (sP93 <=> ((~ (e13 = unit1) & (e13 = op1(e13, e13))) | ~ (e13 = op1(e13, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP93])])).
fof(f118, plain, ((~ (e13 = unit1) & (e14 = op1(e13, e13))) | ~ (e13 = op1(e13, e14)) | ~ sP94), inference(usedef, [], [e118])).
fof(e118, plain, (sP94 <=> ((~ (e13 = unit1) & (e14 = op1(e13, e13))) | ~ (e13 = op1(e13, e14)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP94])])).
fof(f119, plain, ((~ (e13 = unit1) & (e10 = op1(e14, e13))) | ~ (e13 = op1(e14, e10)) | ~ sP95), inference(usedef, [], [e119])).
fof(e119, plain, (sP95 <=> ((~ (e13 = unit1) & (e10 = op1(e14, e13))) | ~ (e13 = op1(e14, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP95])])).
fof(f120, plain, ((~ (e13 = unit1) & (e11 = op1(e14, e13))) | ~ (e13 = op1(e14, e11)) | ~ sP96), inference(usedef, [], [e120])).
fof(e120, plain, (sP96 <=> ((~ (e13 = unit1) & (e11 = op1(e14, e13))) | ~ (e13 = op1(e14, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP96])])).
fof(f121, plain, ((~ (e13 = unit1) & (e12 = op1(e14, e13))) | ~ (e13 = op1(e14, e12)) | ~ sP97), inference(usedef, [], [e121])).
fof(e121, plain, (sP97 <=> ((~ (e13 = unit1) & (e12 = op1(e14, e13))) | ~ (e13 = op1(e14, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP97])])).
fof(f122, plain, ((~ (e13 = unit1) & (e13 = op1(e14, e13))) | ~ (e13 = op1(e14, e13)) | ~ sP98), inference(usedef, [], [e122])).
fof(e122, plain, (sP98 <=> ((~ (e13 = unit1) & (e13 = op1(e14, e13))) | ~ (e13 = op1(e14, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP98])])).
fof(f123, plain, ((~ (e13 = unit1) & (e14 = op1(e14, e13))) | ~ (e13 = op1(e14, e14)) | ~ sP99), inference(usedef, [], [e123])).
fof(e123, plain, (sP99 <=> ((~ (e13 = unit1) & (e14 = op1(e14, e13))) | ~ (e13 = op1(e14, e14)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP99])])).
fof(f124, plain, ((~ (e14 = unit1) & (e10 = op1(e10, e14))) | ~ (op1(e10, e10) = e14) | ~ sP100), inference(usedef, [], [e124])).
fof(e124, plain, (sP100 <=> ((~ (e14 = unit1) & (e10 = op1(e10, e14))) | ~ (op1(e10, e10) = e14))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP100])])).
fof(f125, plain, ((~ (e14 = unit1) & (e11 = op1(e10, e14))) | ~ (e14 = op1(e10, e11)) | ~ sP101), inference(usedef, [], [e125])).
fof(e125, plain, (sP101 <=> ((~ (e14 = unit1) & (e11 = op1(e10, e14))) | ~ (e14 = op1(e10, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP101])])).
fof(f126, plain, ((~ (e14 = unit1) & (e12 = op1(e10, e14))) | ~ (e14 = op1(e10, e12)) | ~ sP102), inference(usedef, [], [e126])).
fof(e126, plain, (sP102 <=> ((~ (e14 = unit1) & (e12 = op1(e10, e14))) | ~ (e14 = op1(e10, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP102])])).
fof(f127, plain, ((~ (e14 = unit1) & (e13 = op1(e10, e14))) | ~ (e14 = op1(e10, e13)) | ~ sP103), inference(usedef, [], [e127])).
fof(e127, plain, (sP103 <=> ((~ (e14 = unit1) & (e13 = op1(e10, e14))) | ~ (e14 = op1(e10, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP103])])).
fof(f128, plain, ((~ (e14 = unit1) & (e14 = op1(e10, e14))) | ~ (e14 = op1(e10, e14)) | ~ sP104), inference(usedef, [], [e128])).
fof(e128, plain, (sP104 <=> ((~ (e14 = unit1) & (e14 = op1(e10, e14))) | ~ (e14 = op1(e10, e14)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP104])])).
fof(f129, plain, ((~ (e14 = unit1) & (e10 = op1(e11, e14))) | ~ (e14 = op1(e11, e10)) | ~ sP105), inference(usedef, [], [e129])).
fof(e129, plain, (sP105 <=> ((~ (e14 = unit1) & (e10 = op1(e11, e14))) | ~ (e14 = op1(e11, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP105])])).
fof(f130, plain, ((~ (e14 = unit1) & (e11 = op1(e11, e14))) | ~ (e14 = op1(e11, e11)) | ~ sP106), inference(usedef, [], [e130])).
fof(e130, plain, (sP106 <=> ((~ (e14 = unit1) & (e11 = op1(e11, e14))) | ~ (e14 = op1(e11, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP106])])).
fof(f131, plain, ((~ (e14 = unit1) & (e12 = op1(e11, e14))) | ~ (e14 = op1(e11, e12)) | ~ sP107), inference(usedef, [], [e131])).
fof(e131, plain, (sP107 <=> ((~ (e14 = unit1) & (e12 = op1(e11, e14))) | ~ (e14 = op1(e11, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP107])])).
fof(f132, plain, ((~ (e14 = unit1) & (e13 = op1(e11, e14))) | ~ (e14 = op1(e11, e13)) | ~ sP108), inference(usedef, [], [e132])).
fof(e132, plain, (sP108 <=> ((~ (e14 = unit1) & (e13 = op1(e11, e14))) | ~ (e14 = op1(e11, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP108])])).
fof(f133, plain, ((~ (e14 = unit1) & (e14 = op1(e11, e14))) | ~ (e14 = op1(e11, e14)) | ~ sP109), inference(usedef, [], [e133])).
fof(e133, plain, (sP109 <=> ((~ (e14 = unit1) & (e14 = op1(e11, e14))) | ~ (e14 = op1(e11, e14)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP109])])).
fof(f134, plain, ((~ (e14 = unit1) & (e10 = op1(e12, e14))) | ~ (e14 = op1(e12, e10)) | ~ sP110), inference(usedef, [], [e134])).
fof(e134, plain, (sP110 <=> ((~ (e14 = unit1) & (e10 = op1(e12, e14))) | ~ (e14 = op1(e12, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP110])])).
fof(f135, plain, ((~ (e14 = unit1) & (e11 = op1(e12, e14))) | ~ (e14 = op1(e12, e11)) | ~ sP111), inference(usedef, [], [e135])).
fof(e135, plain, (sP111 <=> ((~ (e14 = unit1) & (e11 = op1(e12, e14))) | ~ (e14 = op1(e12, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP111])])).
fof(f136, plain, ((~ (e14 = unit1) & (e12 = op1(e12, e14))) | ~ (e14 = op1(e12, e12)) | ~ sP112), inference(usedef, [], [e136])).
fof(e136, plain, (sP112 <=> ((~ (e14 = unit1) & (e12 = op1(e12, e14))) | ~ (e14 = op1(e12, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP112])])).
fof(f137, plain, ((~ (e14 = unit1) & (e13 = op1(e12, e14))) | ~ (e14 = op1(e12, e13)) | ~ sP113), inference(usedef, [], [e137])).
fof(e137, plain, (sP113 <=> ((~ (e14 = unit1) & (e13 = op1(e12, e14))) | ~ (e14 = op1(e12, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP113])])).
fof(f138, plain, ((~ (e14 = unit1) & (e14 = op1(e12, e14))) | ~ (e14 = op1(e12, e14)) | ~ sP114), inference(usedef, [], [e138])).
fof(e138, plain, (sP114 <=> ((~ (e14 = unit1) & (e14 = op1(e12, e14))) | ~ (e14 = op1(e12, e14)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP114])])).
fof(f139, plain, ((~ (e14 = unit1) & (e10 = op1(e13, e14))) | ~ (e14 = op1(e13, e10)) | ~ sP115), inference(usedef, [], [e139])).
fof(e139, plain, (sP115 <=> ((~ (e14 = unit1) & (e10 = op1(e13, e14))) | ~ (e14 = op1(e13, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP115])])).
fof(f140, plain, ((~ (e14 = unit1) & (e11 = op1(e13, e14))) | ~ (e14 = op1(e13, e11)) | ~ sP116), inference(usedef, [], [e140])).
fof(e140, plain, (sP116 <=> ((~ (e14 = unit1) & (e11 = op1(e13, e14))) | ~ (e14 = op1(e13, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP116])])).
fof(f141, plain, ((~ (e14 = unit1) & (e12 = op1(e13, e14))) | ~ (e14 = op1(e13, e12)) | ~ sP117), inference(usedef, [], [e141])).
fof(e141, plain, (sP117 <=> ((~ (e14 = unit1) & (e12 = op1(e13, e14))) | ~ (e14 = op1(e13, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP117])])).
fof(f142, plain, ((~ (e14 = unit1) & (e13 = op1(e13, e14))) | ~ (e14 = op1(e13, e13)) | ~ sP118), inference(usedef, [], [e142])).
fof(e142, plain, (sP118 <=> ((~ (e14 = unit1) & (e13 = op1(e13, e14))) | ~ (e14 = op1(e13, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP118])])).
fof(f143, plain, ((~ (e14 = unit1) & (e14 = op1(e13, e14))) | ~ (e14 = op1(e13, e14)) | ~ sP119), inference(usedef, [], [e143])).
fof(e143, plain, (sP119 <=> ((~ (e14 = unit1) & (e14 = op1(e13, e14))) | ~ (e14 = op1(e13, e14)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP119])])).
fof(f144, plain, ((~ (e14 = unit1) & (e10 = op1(e14, e14))) | ~ (e14 = op1(e14, e10)) | ~ sP120), inference(usedef, [], [e144])).
fof(e144, plain, (sP120 <=> ((~ (e14 = unit1) & (e10 = op1(e14, e14))) | ~ (e14 = op1(e14, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP120])])).
fof(f145, plain, ((~ (e14 = unit1) & (e11 = op1(e14, e14))) | ~ (e14 = op1(e14, e11)) | ~ sP121), inference(usedef, [], [e145])).
fof(e145, plain, (sP121 <=> ((~ (e14 = unit1) & (e11 = op1(e14, e14))) | ~ (e14 = op1(e14, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP121])])).
fof(f146, plain, ((~ (e14 = unit1) & (e12 = op1(e14, e14))) | ~ (e14 = op1(e14, e12)) | ~ sP122), inference(usedef, [], [e146])).
fof(e146, plain, (sP122 <=> ((~ (e14 = unit1) & (e12 = op1(e14, e14))) | ~ (e14 = op1(e14, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP122])])).
fof(f147, plain, ((~ (e14 = unit1) & (e13 = op1(e14, e14))) | ~ (e14 = op1(e14, e13)) | ~ sP123), inference(usedef, [], [e147])).
fof(e147, plain, (sP123 <=> ((~ (e14 = unit1) & (e13 = op1(e14, e14))) | ~ (e14 = op1(e14, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP123])])).
fof(f148, plain, ((~ (e14 = unit1) & (e14 = op1(e14, e14))) | ~ (e14 = op1(e14, e14)) | ~ sP124), inference(usedef, [], [e148])).
fof(e148, plain, (sP124 <=> ((~ (e14 = unit1) & (e14 = op1(e14, e14))) | ~ (e14 = op1(e14, e14)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP124])])).
fof(f149, plain, ((sP24 & sP23 & sP22 & sP21 & sP20 & sP19 & sP18 & sP17 & sP16 & sP15 & sP14 & sP13 & sP12 & sP11 & sP10 & sP9 & sP8 & sP7 & sP6 & sP5 & sP4 & sP3 & sP2 & sP1 & sP0) | ~ sP125), inference(usedef, [], [e149])).
fof(e149, plain, (sP125 <=> (sP24 & sP23 & sP22 & sP21 & sP20 & sP19 & sP18 & sP17 & sP16 & sP15 & sP14 & sP13 & sP12 & sP11 & sP10 & sP9 & sP8 & sP7 & sP6 & sP5 & sP4 & sP3 & sP2 & sP1 & sP0)), introduced(predicate_definition_introduction, [new_symbols(naming, [sP125])])).
fof(f150, plain, ((sP49 & sP48 & sP47 & sP46 & sP45 & sP44 & sP43 & sP42 & sP41 & sP40 & sP39 & sP38 & sP37 & sP36 & sP35 & sP34 & sP33 & sP32 & sP31 & sP30 & sP29 & sP28 & sP27 & sP26 & sP25) | ~ sP126), inference(usedef, [], [e150])).
fof(e150, plain, (sP126 <=> (sP49 & sP48 & sP47 & sP46 & sP45 & sP44 & sP43 & sP42 & sP41 & sP40 & sP39 & sP38 & sP37 & sP36 & sP35 & sP34 & sP33 & sP32 & sP31 & sP30 & sP29 & sP28 & sP27 & sP26 & sP25)), introduced(predicate_definition_introduction, [new_symbols(naming, [sP126])])).
fof(f151, plain, ((sP74 & sP73 & sP72 & sP71 & sP70 & sP69 & sP68 & sP67 & sP66 & sP65 & sP64 & sP63 & sP62 & sP61 & sP60 & sP59 & sP58 & sP57 & sP56 & sP55 & sP54 & sP53 & sP52 & sP51 & sP50) | ~ sP127), inference(usedef, [], [e151])).
fof(e151, plain, (sP127 <=> (sP74 & sP73 & sP72 & sP71 & sP70 & sP69 & sP68 & sP67 & sP66 & sP65 & sP64 & sP63 & sP62 & sP61 & sP60 & sP59 & sP58 & sP57 & sP56 & sP55 & sP54 & sP53 & sP52 & sP51 & sP50)), introduced(predicate_definition_introduction, [new_symbols(naming, [sP127])])).
fof(f152, plain, ((sP99 & sP98 & sP97 & sP96 & sP95 & sP94 & sP93 & sP92 & sP91 & sP90 & sP89 & sP88 & sP87 & sP86 & sP85 & sP84 & sP83 & sP82 & sP81 & sP80 & sP79 & sP78 & sP77 & sP76 & sP75) | ~ sP128), inference(usedef, [], [e152])).
fof(e152, plain, (sP128 <=> (sP99 & sP98 & sP97 & sP96 & sP95 & sP94 & sP93 & sP92 & sP91 & sP90 & sP89 & sP88 & sP87 & sP86 & sP85 & sP84 & sP83 & sP82 & sP81 & sP80 & sP79 & sP78 & sP77 & sP76 & sP75)), introduced(predicate_definition_introduction, [new_symbols(naming, [sP128])])).
fof(f153, plain, ((sP124 & sP123 & sP122 & sP121 & sP120 & sP119 & sP118 & sP117 & sP116 & sP115 & sP114 & sP113 & sP112 & sP111 & sP110 & sP109 & sP108 & sP107 & sP106 & sP105 & sP104 & sP103 & sP102 & sP101 & sP100) | ~ sP129), inference(usedef, [], [e153])).
fof(e153, plain, (sP129 <=> (sP124 & sP123 & sP122 & sP121 & sP120 & sP119 & sP118 & sP117 & sP116 & sP115 & sP114 & sP113 & sP112 & sP111 & sP110 & sP109 & sP108 & sP107 & sP106 & sP105 & sP104 & sP103 & sP102 & sP101 & sP100)), introduced(predicate_definition_introduction, [new_symbols(naming, [sP129])])).
fof(f12, plain, ((((~ (e14 = unit1) & (e14 = op1(e14, e14))) | ~ (e14 = op1(e14, e14))) & ((~ (e14 = unit1) & (e13 = op1(e14, e14))) | ~ (e14 = op1(e14, e13))) & ((~ (e14 = unit1) & (e12 = op1(e14, e14))) | ~ (e14 = op1(e14, e12))) & ((~ (e14 = unit1) & (e11 = op1(e14, e14))) | ~ (e14 = op1(e14, e11))) & ((~ (e14 = unit1) & (e10 = op1(e14, e14))) | ~ (e14 = op1(e14, e10))) & ((~ (e14 = unit1) & (e14 = op1(e13, e14))) | ~ (e14 = op1(e13, e14))) & ((~ (e14 = unit1) & (e13 = op1(e13, e14))) | ~ (e14 = op1(e13, e13))) & ((~ (e14 = unit1) & (e12 = op1(e13, e14))) | ~ (e14 = op1(e13, e12))) & ((~ (e14 = unit1) & (e11 = op1(e13, e14))) | ~ (e14 = op1(e13, e11))) & ((~ (e14 = unit1) & (e10 = op1(e13, e14))) | ~ (e14 = op1(e13, e10))) & ((~ (e14 = unit1) & (e14 = op1(e12, e14))) | ~ (e14 = op1(e12, e14))) & ((~ (e14 = unit1) & (e13 = op1(e12, e14))) | ~ (e14 = op1(e12, e13))) & ((~ (e14 = unit1) & (e12 = op1(e12, e14))) | ~ (e14 = op1(e12, e12))) & ((~ (e14 = unit1) & (e11 = op1(e12, e14))) | ~ (e14 = op1(e12, e11))) & ((~ (e14 = unit1) & (e10 = op1(e12, e14))) | ~ (e14 = op1(e12, e10))) & ((~ (e14 = unit1) & (e14 = op1(e11, e14))) | ~ (e14 = op1(e11, e14))) & ((~ (e14 = unit1) & (e13 = op1(e11, e14))) | ~ (e14 = op1(e11, e13))) & ((~ (e14 = unit1) & (e12 = op1(e11, e14))) | ~ (e14 = op1(e11, e12))) & ((~ (e14 = unit1) & (e11 = op1(e11, e14))) | ~ (e14 = op1(e11, e11))) & ((~ (e14 = unit1) & (e10 = op1(e11, e14))) | ~ (e14 = op1(e11, e10))) & ((~ (e14 = unit1) & (e14 = op1(e10, e14))) | ~ (e14 = op1(e10, e14))) & ((~ (e14 = unit1) & (e13 = op1(e10, e14))) | ~ (e14 = op1(e10, e13))) & ((~ (e14 = unit1) & (e12 = op1(e10, e14))) | ~ (e14 = op1(e10, e12))) & ((~ (e14 = unit1) & (e11 = op1(e10, e14))) | ~ (e14 = op1(e10, e11))) & ((~ (e14 = unit1) & (e10 = op1(e10, e14))) | ~ (op1(e10, e10) = e14))) | (((~ (e13 = unit1) & (e14 = op1(e14, e13))) | ~ (e13 = op1(e14, e14))) & ((~ (e13 = unit1) & (e13 = op1(e14, e13))) | ~ (e13 = op1(e14, e13))) & ((~ (e13 = unit1) & (e12 = op1(e14, e13))) | ~ (e13 = op1(e14, e12))) & ((~ (e13 = unit1) & (e11 = op1(e14, e13))) | ~ (e13 = op1(e14, e11))) & ((~ (e13 = unit1) & (e10 = op1(e14, e13))) | ~ (e13 = op1(e14, e10))) & ((~ (e13 = unit1) & (e14 = op1(e13, e13))) | ~ (e13 = op1(e13, e14))) & ((~ (e13 = unit1) & (e13 = op1(e13, e13))) | ~ (e13 = op1(e13, e13))) & ((~ (e13 = unit1) & (e12 = op1(e13, e13))) | ~ (e13 = op1(e13, e12))) & ((~ (e13 = unit1) & (e11 = op1(e13, e13))) | ~ (e13 = op1(e13, e11))) & ((~ (e13 = unit1) & (e10 = op1(e13, e13))) | ~ (e13 = op1(e13, e10))) & ((~ (e13 = unit1) & (e14 = op1(e12, e13))) | ~ (e13 = op1(e12, e14))) & ((~ (e13 = unit1) & (e13 = op1(e12, e13))) | ~ (e13 = op1(e12, e13))) & ((~ (e13 = unit1) & (e12 = op1(e12, e13))) | ~ (e13 = op1(e12, e12))) & ((~ (e13 = unit1) & (e11 = op1(e12, e13))) | ~ (e13 = op1(e12, e11))) & ((~ (e13 = unit1) & (e10 = op1(e12, e13))) | ~ (e13 = op1(e12, e10))) & ((~ (e13 = unit1) & (e14 = op1(e11, e13))) | ~ (e13 = op1(e11, e14))) & ((~ (e13 = unit1) & (e13 = op1(e11, e13))) | ~ (e13 = op1(e11, e13))) & ((~ (e13 = unit1) & (e12 = op1(e11, e13))) | ~ (e13 = op1(e11, e12))) & ((~ (e13 = unit1) & (e11 = op1(e11, e13))) | ~ (e13 = op1(e11, e11))) & ((~ (e13 = unit1) & (e10 = op1(e11, e13))) | ~ (e13 = op1(e11, e10))) & ((~ (e13 = unit1) & (e14 = op1(e10, e13))) | ~ (e13 = op1(e10, e14))) & ((~ (e13 = unit1) & (e13 = op1(e10, e13))) | ~ (e13 = op1(e10, e13))) & ((~ (e13 = unit1) & (e12 = op1(e10, e13))) | ~ (e13 = op1(e10, e12))) & ((~ (e13 = unit1) & (e11 = op1(e10, e13))) | ~ (e13 = op1(e10, e11))) & ((~ (e13 = unit1) & (e10 = op1(e10, e13))) | ~ (op1(e10, e10) = e13))) | (((~ (e12 = unit1) & (e14 = op1(e14, e12))) | ~ (e12 = op1(e14, e14))) & ((~ (e12 = unit1) & (e13 = op1(e14, e12))) | ~ (e12 = op1(e14, e13))) & ((~ (e12 = unit1) & (e12 = op1(e14, e12))) | ~ (e12 = op1(e14, e12))) & ((~ (e12 = unit1) & (e11 = op1(e14, e12))) | ~ (e12 = op1(e14, e11))) & ((~ (e12 = unit1) & (e10 = op1(e14, e12))) | ~ (e12 = op1(e14, e10))) & ((~ (e12 = unit1) & (e14 = op1(e13, e12))) | ~ (e12 = op1(e13, e14))) & ((~ (e12 = unit1) & (e13 = op1(e13, e12))) | ~ (e12 = op1(e13, e13))) & ((~ (e12 = unit1) & (e12 = op1(e13, e12))) | ~ (e12 = op1(e13, e12))) & ((~ (e12 = unit1) & (e11 = op1(e13, e12))) | ~ (e12 = op1(e13, e11))) & ((~ (e12 = unit1) & (e10 = op1(e13, e12))) | ~ (e12 = op1(e13, e10))) & ((~ (e12 = unit1) & (e14 = op1(e12, e12))) | ~ (e12 = op1(e12, e14))) & ((~ (e12 = unit1) & (e13 = op1(e12, e12))) | ~ (e12 = op1(e12, e13))) & ((~ (e12 = unit1) & (e12 = op1(e12, e12))) | ~ (e12 = op1(e12, e12))) & ((~ (e12 = unit1) & (e11 = op1(e12, e12))) | ~ (e12 = op1(e12, e11))) & ((~ (e12 = unit1) & (e10 = op1(e12, e12))) | ~ (e12 = op1(e12, e10))) & ((~ (e12 = unit1) & (e14 = op1(e11, e12))) | ~ (e12 = op1(e11, e14))) & ((~ (e12 = unit1) & (e13 = op1(e11, e12))) | ~ (e12 = op1(e11, e13))) & ((~ (e12 = unit1) & (e12 = op1(e11, e12))) | ~ (e12 = op1(e11, e12))) & ((~ (e12 = unit1) & (e11 = op1(e11, e12))) | ~ (e12 = op1(e11, e11))) & ((~ (e12 = unit1) & (e10 = op1(e11, e12))) | ~ (e12 = op1(e11, e10))) & ((~ (e12 = unit1) & (e14 = op1(e10, e12))) | ~ (e12 = op1(e10, e14))) & ((~ (e12 = unit1) & (e13 = op1(e10, e12))) | ~ (e12 = op1(e10, e13))) & ((~ (e12 = unit1) & (e12 = op1(e10, e12))) | ~ (e12 = op1(e10, e12))) & ((~ (e12 = unit1) & (e11 = op1(e10, e12))) | ~ (e12 = op1(e10, e11))) & ((~ (e12 = unit1) & (e10 = op1(e10, e12))) | ~ (op1(e10, e10) = e12))) | (((~ (e11 = unit1) & (e14 = op1(e14, e11))) | ~ (e11 = op1(e14, e14))) & ((~ (e11 = unit1) & (e13 = op1(e14, e11))) | ~ (e11 = op1(e14, e13))) & ((~ (e11 = unit1) & (e12 = op1(e14, e11))) | ~ (e11 = op1(e14, e12))) & ((~ (e11 = unit1) & (e11 = op1(e14, e11))) | ~ (e11 = op1(e14, e11))) & ((~ (e11 = unit1) & (e10 = op1(e14, e11))) | ~ (e11 = op1(e14, e10))) & ((~ (e11 = unit1) & (e14 = op1(e13, e11))) | ~ (e11 = op1(e13, e14))) & ((~ (e11 = unit1) & (e13 = op1(e13, e11))) | ~ (e11 = op1(e13, e13))) & ((~ (e11 = unit1) & (e12 = op1(e13, e11))) | ~ (e11 = op1(e13, e12))) & ((~ (e11 = unit1) & (e11 = op1(e13, e11))) | ~ (e11 = op1(e13, e11))) & ((~ (e11 = unit1) & (e10 = op1(e13, e11))) | ~ (e11 = op1(e13, e10))) & ((~ (e11 = unit1) & (e14 = op1(e12, e11))) | ~ (e11 = op1(e12, e14))) & ((~ (e11 = unit1) & (e13 = op1(e12, e11))) | ~ (e11 = op1(e12, e13))) & ((~ (e11 = unit1) & (e12 = op1(e12, e11))) | ~ (e11 = op1(e12, e12))) & ((~ (e11 = unit1) & (e11 = op1(e12, e11))) | ~ (e11 = op1(e12, e11))) & ((~ (e11 = unit1) & (e10 = op1(e12, e11))) | ~ (e11 = op1(e12, e10))) & ((~ (e11 = unit1) & (e14 = op1(e11, e11))) | ~ (e11 = op1(e11, e14))) & ((~ (e11 = unit1) & (e13 = op1(e11, e11))) | ~ (e11 = op1(e11, e13))) & ((~ (e11 = unit1) & (e12 = op1(e11, e11))) | ~ (e11 = op1(e11, e12))) & ((~ (e11 = unit1) & (e11 = op1(e11, e11))) | ~ (e11 = op1(e11, e11))) & ((~ (e11 = unit1) & (e10 = op1(e11, e11))) | ~ (e11 = op1(e11, e10))) & ((~ (e11 = unit1) & (e14 = op1(e10, e11))) | ~ (e11 = op1(e10, e14))) & ((~ (e11 = unit1) & (e13 = op1(e10, e11))) | ~ (e11 = op1(e10, e13))) & ((~ (e11 = unit1) & (e12 = op1(e10, e11))) | ~ (e11 = op1(e10, e12))) & ((~ (e11 = unit1) & (e11 = op1(e10, e11))) | ~ (e11 = op1(e10, e11))) & ((~ (e11 = unit1) & (e10 = op1(e10, e11))) | ~ (op1(e10, e10) = e11))) | (((~ (e10 = unit1) & (e14 = op1(e14, e10))) | ~ (e10 = op1(e14, e14))) & ((~ (e10 = unit1) & (e13 = op1(e14, e10))) | ~ (e10 = op1(e14, e13))) & ((~ (e10 = unit1) & (e12 = op1(e14, e10))) | ~ (e10 = op1(e14, e12))) & ((~ (e10 = unit1) & (e11 = op1(e14, e10))) | ~ (e10 = op1(e14, e11))) & ((~ (e10 = unit1) & (e10 = op1(e14, e10))) | ~ (e10 = op1(e14, e10))) & ((~ (e10 = unit1) & (e14 = op1(e13, e10))) | ~ (e10 = op1(e13, e14))) & ((~ (e10 = unit1) & (e13 = op1(e13, e10))) | ~ (e10 = op1(e13, e13))) & ((~ (e10 = unit1) & (e12 = op1(e13, e10))) | ~ (e10 = op1(e13, e12))) & ((~ (e10 = unit1) & (e11 = op1(e13, e10))) | ~ (e10 = op1(e13, e11))) & ((~ (e10 = unit1) & (e10 = op1(e13, e10))) | ~ (e10 = op1(e13, e10))) & ((~ (e10 = unit1) & (e14 = op1(e12, e10))) | ~ (e10 = op1(e12, e14))) & ((~ (e10 = unit1) & (e13 = op1(e12, e10))) | ~ (e10 = op1(e12, e13))) & ((~ (e10 = unit1) & (e12 = op1(e12, e10))) | ~ (e10 = op1(e12, e12))) & ((~ (e10 = unit1) & (e11 = op1(e12, e10))) | ~ (e10 = op1(e12, e11))) & ((~ (e10 = unit1) & (e10 = op1(e12, e10))) | ~ (e10 = op1(e12, e10))) & ((~ (e10 = unit1) & (e14 = op1(e11, e10))) | ~ (e10 = op1(e11, e14))) & ((~ (e10 = unit1) & (e13 = op1(e11, e10))) | ~ (e10 = op1(e11, e13))) & ((~ (e10 = unit1) & (e12 = op1(e11, e10))) | ~ (e10 = op1(e11, e12))) & ((~ (e10 = unit1) & (e11 = op1(e11, e10))) | ~ (e10 = op1(e11, e11))) & ((~ (e10 = unit1) & (e10 = op1(e11, e10))) | ~ (e10 = op1(e11, e10))) & ((~ (e10 = unit1) & (op1(e10, e10) = e14)) | ~ (e10 = op1(e10, e14))) & ((~ (e10 = unit1) & (op1(e10, e10) = e13)) | ~ (e10 = op1(e10, e13))) & ((~ (e10 = unit1) & (op1(e10, e10) = e12)) | ~ (e10 = op1(e10, e12))) & ((~ (e10 = unit1) & (op1(e10, e10) = e11)) | ~ (e10 = op1(e10, e11))) & ((~ (e10 = unit1) & (e10 = op1(e10, e10))) | ~ (e10 = op1(e10, e10))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG051+1.p', ax12)).
fof(f3964, plain, (~ spl280_384 | ~ spl280_91 | ~ spl280_126), inference(avatar_split_clause, [], [f1366, f2440, f2293, f3717])).
fof(f3717, plain, (spl280_384 <=> sP6), introduced(avatar_definition, [new_symbols(naming, [spl280_384])])).
fof(f1366, plain, (~ (e10 = unit1) | ~ (e10 = op1(e11, e11)) | ~ sP6), inference(cnf_transformation, [], [f430])).
fof(f430, plain, ((~ (e10 = unit1) & (e11 = op1(e11, e10))) | ~ (e10 = op1(e11, e11)) | ~ sP6), inference(nnf_transformation, [], [f30])).
fof(f3923, plain, (~ spl280_360 | ~ spl280_102 | spl280_120), inference(avatar_split_clause, [], [f1319, f2414, f2339, f3598])).
fof(f3598, plain, (spl280_360 <=> sP29), introduced(avatar_definition, [new_symbols(naming, [spl280_360])])).
fof(f2339, plain, (spl280_102 <=> (e11 = op1(e10, e14))), introduced(avatar_definition, [new_symbols(naming, [spl280_102])])).
fof(f1319, plain, ((e14 = op1(e10, e11)) | ~ (e11 = op1(e10, e14)) | ~ sP29), inference(cnf_transformation, [], [f407])).
fof(f407, plain, ((~ (e11 = unit1) & (e14 = op1(e10, e11))) | ~ (e11 = op1(e10, e14)) | ~ sP29), inference(nnf_transformation, [], [f53])).
fof(f3900, plain, (~ spl280_347 | ~ spl280_37 | spl280_43), inference(avatar_split_clause, [], [f1293, f2091, f2066, f3533])).
fof(f3533, plain, (spl280_347 <=> sP42), introduced(avatar_definition, [new_symbols(naming, [spl280_347])])).
fof(f1293, plain, ((e12 = op1(e13, e11)) | ~ (e11 = op1(e13, e12)) | ~ sP42), inference(cnf_transformation, [], [f394])).
fof(f394, plain, ((~ (e11 = unit1) & (e12 = op1(e13, e11))) | ~ (e11 = op1(e13, e12)) | ~ sP42), inference(nnf_transformation, [], [f66])).
fof(f3887, plain, (~ spl280_340 | ~ spl280_2 | spl280_20), inference(avatar_split_clause, [], [f1279, f1994, f1919, f3498])).
fof(f3498, plain, (spl280_340 <=> sP49), introduced(avatar_definition, [new_symbols(naming, [spl280_340])])).
fof(f1919, plain, (spl280_2 <=> (e11 = op1(e14, e14))), introduced(avatar_definition, [new_symbols(naming, [spl280_2])])).
fof(f1279, plain, ((e14 = op1(e14, e11)) | ~ (e11 = op1(e14, e14)) | ~ sP49), inference(cnf_transformation, [], [f387])).
fof(f387, plain, ((~ (e11 = unit1) & (e14 = op1(e14, e11))) | ~ (e11 = op1(e14, e14)) | ~ sP49), inference(nnf_transformation, [], [f73])).
fof(f3867, plain, (~ spl280_328 | ~ spl280_73 | spl280_61), inference(avatar_split_clause, [], [f1257, f2167, f2217, f3439])).
fof(f3439, plain, (spl280_328 <=> sP60), introduced(avatar_definition, [new_symbols(naming, [spl280_328])])).
fof(f1257, plain, ((e10 = op1(e12, e12)) | ~ (e12 = op1(e12, e10)) | ~ sP60), inference(cnf_transformation, [], [f376])).
fof(f376, plain, ((~ (e12 = unit1) & (e10 = op1(e12, e12))) | ~ (e12 = op1(e12, e10)) | ~ sP60), inference(nnf_transformation, [], [f84])).
fof(f3813, plain, (~ spl280_297 | ~ spl280_49 | spl280_31), inference(avatar_split_clause, [], [f1197, f2041, f2116, f3285])).
fof(f3285, plain, (spl280_297 <=> sP90), introduced(avatar_definition, [new_symbols(naming, [spl280_297])])).
fof(f1197, plain, ((e10 = op1(e13, e13)) | ~ (e13 = op1(e13, e10)) | ~ sP90), inference(cnf_transformation, [], [f346])).
fof(f346, plain, ((~ (e13 = unit1) & (e10 = op1(e13, e13))) | ~ (e13 = op1(e13, e10)) | ~ sP90), inference(nnf_transformation, [], [f114])).
fof(f3762, plain, (~ spl280_268 | ~ spl280_35 | spl280_29), inference(avatar_split_clause, [], [f1141, f2032, f2057, f3141])).
fof(f3141, plain, (spl280_268 <=> sP118), introduced(avatar_definition, [new_symbols(naming, [spl280_268])])).
fof(f1141, plain, ((e13 = op1(e13, e14)) | ~ (e14 = op1(e13, e13)) | ~ sP118), inference(cnf_transformation, [], [f318])).
fof(f318, plain, ((~ (e14 = unit1) & (e13 = op1(e13, e14))) | ~ (e14 = op1(e13, e13)) | ~ sP118), inference(nnf_transformation, [], [f142])).
fof(f3720, plain, (~ spl280_365 | spl280_384), inference(avatar_split_clause, [], [f1110, f3717, f3623])).
fof(f1110, plain, (sP6 | ~ sP125), inference(cnf_transformation, [], [f311])).
fof(f311, plain, ((sP24 & sP23 & sP22 & sP21 & sP20 & sP19 & sP18 & sP17 & sP16 & sP15 & sP14 & sP13 & sP12 & sP11 & sP10 & sP9 & sP8 & sP7 & sP6 & sP5 & sP4 & sP3 & sP2 & sP1 & sP0) | ~ sP125), inference(nnf_transformation, [], [f149])).
fof(f3601, plain, (~ spl280_339 | spl280_360), inference(avatar_split_clause, [], [f1083, f3598, f3494])).
fof(f1083, plain, (sP29 | ~ sP126), inference(cnf_transformation, [], [f310])).
fof(f310, plain, ((sP49 & sP48 & sP47 & sP46 & sP45 & sP44 & sP43 & sP42 & sP41 & sP40 & sP39 & sP38 & sP37 & sP36 & sP35 & sP34 & sP33 & sP32 & sP31 & sP30 & sP29 & sP28 & sP27 & sP26 & sP25) | ~ sP126), inference(nnf_transformation, [], [f150])).
fof(f3536, plain, (~ spl280_339 | spl280_347), inference(avatar_split_clause, [], [f1096, f3533, f3494])).
fof(f1096, plain, (sP42 | ~ sP126), inference(cnf_transformation, [], [f310])).
fof(f3501, plain, (~ spl280_339 | spl280_340), inference(avatar_split_clause, [], [f1103, f3498, f3494])).
fof(f1103, plain, (sP49 | ~ sP126), inference(cnf_transformation, [], [f310])).
fof(f3442, plain, (~ spl280_313 | spl280_328), inference(avatar_split_clause, [], [f1064, f3439, f3365])).
fof(f1064, plain, (sP60 | ~ sP127), inference(cnf_transformation, [], [f309])).
fof(f309, plain, ((sP74 & sP73 & sP72 & sP71 & sP70 & sP69 & sP68 & sP67 & sP66 & sP65 & sP64 & sP63 & sP62 & sP61 & sP60 & sP59 & sP58 & sP57 & sP56 & sP55 & sP54 & sP53 & sP52 & sP51 & sP50) | ~ sP127), inference(nnf_transformation, [], [f151])).
fof(f3288, plain, (~ spl280_287 | spl280_297), inference(avatar_split_clause, [], [f1044, f3285, f3236])).
fof(f1044, plain, (sP90 | ~ sP128), inference(cnf_transformation, [], [f308])).
fof(f308, plain, ((sP99 & sP98 & sP97 & sP96 & sP95 & sP94 & sP93 & sP92 & sP91 & sP90 & sP89 & sP88 & sP87 & sP86 & sP85 & sP84 & sP83 & sP82 & sP81 & sP80 & sP79 & sP78 & sP77 & sP76 & sP75) | ~ sP128), inference(nnf_transformation, [], [f152])).
fof(f3144, plain, (~ spl280_261 | spl280_268), inference(avatar_split_clause, [], [f1022, f3141, f3107])).
fof(f1022, plain, (sP118 | ~ sP129), inference(cnf_transformation, [], [f307])).
fof(f307, plain, ((sP124 & sP123 & sP122 & sP121 & sP120 & sP119 & sP118 & sP117 & sP116 & sP115 & sP114 & sP113 & sP112 & sP111 & sP110 & sP109 & sP108 & sP107 & sP106 & sP105 & sP104 & sP103 & sP102 & sP101 & sP100) | ~ sP129), inference(nnf_transformation, [], [f153])).
fof(f3089, plain, (spl280_229 | spl280_224 | spl280_219 | spl280_214 | spl280_209), inference(avatar_split_clause, [], [f725, f2838, f2859, f2880, f2901, f2922])).
fof(f725, plain, ((e23 = op2(e21, e24)) | (e23 = op2(e21, e23)) | (e23 = op2(e21, e22)) | (e23 = op2(e21, e21)) | (e23 = op2(e21, e20))), inference(cnf_transformation, [], [f6])).
fof(f6, plain, (((e24 = op2(e24, e24)) | (e24 = op2(e23, e24)) | (e24 = op2(e22, e24)) | (e24 = op2(e21, e24)) | (e24 = op2(e20, e24))) & ((e24 = op2(e24, e24)) | (e24 = op2(e24, e23)) | (e24 = op2(e24, e22)) | (e24 = op2(e24, e21)) | (e24 = op2(e24, e20))) & ((e23 = op2(e24, e24)) | (e23 = op2(e23, e24)) | (e23 = op2(e22, e24)) | (e23 = op2(e21, e24)) | (e23 = op2(e20, e24))) & ((e23 = op2(e24, e24)) | (e23 = op2(e24, e23)) | (e23 = op2(e24, e22)) | (e23 = op2(e24, e21)) | (e23 = op2(e24, e20))) & ((e22 = op2(e24, e24)) | (e22 = op2(e23, e24)) | (e22 = op2(e22, e24)) | (e22 = op2(e21, e24)) | (e22 = op2(e20, e24))) & ((e22 = op2(e24, e24)) | (e22 = op2(e24, e23)) | (e22 = op2(e24, e22)) | (e22 = op2(e24, e21)) | (e22 = op2(e24, e20))) & ((e21 = op2(e24, e24)) | (e21 = op2(e23, e24)) | (e21 = op2(e22, e24)) | (e21 = op2(e21, e24)) | (e21 = op2(e20, e24))) & ((e21 = op2(e24, e24)) | (e21 = op2(e24, e23)) | (e21 = op2(e24, e22)) | (e21 = op2(e24, e21)) | (e21 = op2(e24, e20))) & ((e20 = op2(e24, e24)) | (e20 = op2(e23, e24)) | (e20 = op2(e22, e24)) | (e20 = op2(e21, e24)) | (e20 = op2(e20, e24))) & ((e20 = op2(e24, e24)) | (e20 = op2(e24, e23)) | (e20 = op2(e24, e22)) | (e20 = op2(e24, e21)) | (e20 = op2(e24, e20))) & ((e24 = op2(e24, e23)) | (e24 = op2(e23, e23)) | (e24 = op2(e22, e23)) | (e24 = op2(e21, e23)) | (e24 = op2(e20, e23))) & ((e24 = op2(e23, e24)) | (e24 = op2(e23, e23)) | (e24 = op2(e23, e22)) | (e24 = op2(e23, e21)) | (e24 = op2(e23, e20))) & ((e23 = op2(e24, e23)) | (e23 = op2(e23, e23)) | (e23 = op2(e22, e23)) | (e23 = op2(e21, e23)) | (e23 = op2(e20, e23))) & ((e23 = op2(e23, e24)) | (e23 = op2(e23, e23)) | (e23 = op2(e23, e22)) | (e23 = op2(e23, e21)) | (e23 = op2(e23, e20))) & ((e22 = op2(e24, e23)) | (e22 = op2(e23, e23)) | (e22 = op2(e22, e23)) | (e22 = op2(e21, e23)) | (e22 = op2(e20, e23))) & ((e22 = op2(e23, e24)) | (e22 = op2(e23, e23)) | (e22 = op2(e23, e22)) | (e22 = op2(e23, e21)) | (e22 = op2(e23, e20))) & ((e21 = op2(e24, e23)) | (e21 = op2(e23, e23)) | (e21 = op2(e22, e23)) | (e21 = op2(e21, e23)) | (e21 = op2(e20, e23))) & ((e21 = op2(e23, e24)) | (e21 = op2(e23, e23)) | (e21 = op2(e23, e22)) | (e21 = op2(e23, e21)) | (e21 = op2(e23, e20))) & ((e20 = op2(e24, e23)) | (e20 = op2(e23, e23)) | (e20 = op2(e22, e23)) | (e20 = op2(e21, e23)) | (e20 = op2(e20, e23))) & ((e20 = op2(e23, e24)) | (e20 = op2(e23, e23)) | (e20 = op2(e23, e22)) | (e20 = op2(e23, e21)) | (e20 = op2(e23, e20))) & ((e24 = op2(e24, e22)) | (e24 = op2(e23, e22)) | (e24 = op2(e22, e22)) | (e24 = op2(e21, e22)) | (e24 = op2(e20, e22))) & ((e24 = op2(e22, e24)) | (e24 = op2(e22, e23)) | (e24 = op2(e22, e22)) | (e24 = op2(e22, e21)) | (e24 = op2(e22, e20))) & ((e23 = op2(e24, e22)) | (e23 = op2(e23, e22)) | (e23 = op2(e22, e22)) | (e23 = op2(e21, e22)) | (e23 = op2(e20, e22))) & ((e23 = op2(e22, e24)) | (e23 = op2(e22, e23)) | (e23 = op2(e22, e22)) | (e23 = op2(e22, e21)) | (e23 = op2(e22, e20))) & ((e22 = op2(e24, e22)) | (e22 = op2(e23, e22)) | (e22 = op2(e22, e22)) | (e22 = op2(e21, e22)) | (e22 = op2(e20, e22))) & ((e22 = op2(e22, e24)) | (e22 = op2(e22, e23)) | (e22 = op2(e22, e22)) | (e22 = op2(e22, e21)) | (e22 = op2(e22, e20))) & ((e21 = op2(e24, e22)) | (e21 = op2(e23, e22)) | (e21 = op2(e22, e22)) | (e21 = op2(e21, e22)) | (e21 = op2(e20, e22))) & ((e21 = op2(e22, e24)) | (e21 = op2(e22, e23)) | (e21 = op2(e22, e22)) | (e21 = op2(e22, e21)) | (e21 = op2(e22, e20))) & ((e20 = op2(e24, e22)) | (e20 = op2(e23, e22)) | (e20 = op2(e22, e22)) | (e20 = op2(e21, e22)) | (e20 = op2(e20, e22))) & ((e20 = op2(e22, e24)) | (e20 = op2(e22, e23)) | (e20 = op2(e22, e22)) | (e20 = op2(e22, e21)) | (e20 = op2(e22, e20))) & ((e24 = op2(e24, e21)) | (e24 = op2(e23, e21)) | (e24 = op2(e22, e21)) | (e24 = op2(e21, e21)) | (e24 = op2(e20, e21))) & ((e24 = op2(e21, e24)) | (e24 = op2(e21, e23)) | (e24 = op2(e21, e22)) | (e24 = op2(e21, e21)) | (e24 = op2(e21, e20))) & ((e23 = op2(e24, e21)) | (e23 = op2(e23, e21)) | (e23 = op2(e22, e21)) | (e23 = op2(e21, e21)) | (e23 = op2(e20, e21))) & ((e23 = op2(e21, e24)) | (e23 = op2(e21, e23)) | (e23 = op2(e21, e22)) | (e23 = op2(e21, e21)) | (e23 = op2(e21, e20))) & ((e22 = op2(e24, e21)) | (e22 = op2(e23, e21)) | (e22 = op2(e22, e21)) | (e22 = op2(e21, e21)) | (e22 = op2(e20, e21))) & ((e22 = op2(e21, e24)) | (e22 = op2(e21, e23)) | (e22 = op2(e21, e22)) | (e22 = op2(e21, e21)) | (e22 = op2(e21, e20))) & ((e21 = op2(e24, e21)) | (e21 = op2(e23, e21)) | (e21 = op2(e22, e21)) | (e21 = op2(e21, e21)) | (e21 = op2(e20, e21))) & ((e21 = op2(e21, e24)) | (e21 = op2(e21, e23)) | (e21 = op2(e21, e22)) | (e21 = op2(e21, e21)) | (e21 = op2(e21, e20))) & ((e20 = op2(e24, e21)) | (e20 = op2(e23, e21)) | (e20 = op2(e22, e21)) | (e20 = op2(e21, e21)) | (e20 = op2(e20, e21))) & ((e20 = op2(e21, e24)) | (e20 = op2(e21, e23)) | (e20 = op2(e21, e22)) | (e20 = op2(e21, e21)) | (e20 = op2(e21, e20))) & ((e24 = op2(e24, e20)) | (e24 = op2(e23, e20)) | (e24 = op2(e22, e20)) | (e24 = op2(e21, e20)) | (op2(e20, e20) = e24)) & ((e24 = op2(e20, e24)) | (e24 = op2(e20, e23)) | (e24 = op2(e20, e22)) | (e24 = op2(e20, e21)) | (op2(e20, e20) = e24)) & ((e23 = op2(e24, e20)) | (e23 = op2(e23, e20)) | (e23 = op2(e22, e20)) | (e23 = op2(e21, e20)) | (op2(e20, e20) = e23)) & ((e23 = op2(e20, e24)) | (e23 = op2(e20, e23)) | (e23 = op2(e20, e22)) | (e23 = op2(e20, e21)) | (op2(e20, e20) = e23)) & ((e22 = op2(e24, e20)) | (e22 = op2(e23, e20)) | (e22 = op2(e22, e20)) | (e22 = op2(e21, e20)) | (op2(e20, e20) = e22)) & ((e22 = op2(e20, e24)) | (e22 = op2(e20, e23)) | (e22 = op2(e20, e22)) | (e22 = op2(e20, e21)) | (op2(e20, e20) = e22)) & ((e21 = op2(e24, e20)) | (e21 = op2(e23, e20)) | (e21 = op2(e22, e20)) | (e21 = op2(e21, e20)) | (op2(e20, e20) = e21)) & ((e21 = op2(e20, e24)) | (e21 = op2(e20, e23)) | (e21 = op2(e20, e22)) | (e21 = op2(e20, e21)) | (op2(e20, e20) = e21)) & ((e20 = op2(e24, e20)) | (e20 = op2(e23, e20)) | (e20 = op2(e22, e20)) | (e20 = op2(e21, e20)) | (e20 = op2(e20, e20))) & ((e20 = op2(e20, e24)) | (e20 = op2(e20, e23)) | (e20 = op2(e20, e22)) | (e20 = op2(e20, e21)) | (e20 = op2(e20, e20)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG051+1.p', ax6)).
fof(f3087, plain, (spl280_230 | spl280_225 | spl280_220 | spl280_215 | spl280_210), inference(avatar_split_clause, [], [f727, f2842, f2863, f2884, f2905, f2926])).
fof(f727, plain, ((e24 = op2(e21, e24)) | (e24 = op2(e21, e23)) | (e24 = op2(e21, e22)) | (e24 = op2(e21, e21)) | (e24 = op2(e21, e20))), inference(cnf_transformation, [], [f6])).
fof(f3086, plain, (spl280_250 | spl280_225 | spl280_200 | spl280_175 | spl280_150), inference(avatar_split_clause, [], [f728, f2590, f2695, f2800, f2905, f3010])).
fof(f728, plain, ((e24 = op2(e24, e21)) | (e24 = op2(e23, e21)) | (e24 = op2(e22, e21)) | (e24 = op2(e21, e21)) | (e24 = op2(e20, e21))), inference(cnf_transformation, [], [f6])).
fof(f3085, plain, (spl280_201 | spl280_196 | spl280_191 | spl280_186 | spl280_181), inference(avatar_split_clause, [], [f729, f2721, f2742, f2763, f2784, f2805])).
fof(f729, plain, ((e20 = op2(e22, e24)) | (e20 = op2(e22, e23)) | (e20 = op2(e22, e22)) | (e20 = op2(e22, e21)) | (e20 = op2(e22, e20))), inference(cnf_transformation, [], [f6])).
fof(f3084, plain, (spl280_241 | spl280_216 | spl280_191 | spl280_166 | spl280_141), inference(avatar_split_clause, [], [f730, f2553, f2658, f2763, f2868, f2973])).
fof(f730, plain, ((e20 = op2(e24, e22)) | (e20 = op2(e23, e22)) | (e20 = op2(e22, e22)) | (e20 = op2(e21, e22)) | (e20 = op2(e20, e22))), inference(cnf_transformation, [], [f6])).
fof(f3083, plain, (spl280_202 | spl280_197 | spl280_192 | spl280_187 | spl280_182), inference(avatar_split_clause, [], [f731, f2725, f2746, f2767, f2788, f2809])).
fof(f731, plain, ((e21 = op2(e22, e24)) | (e21 = op2(e22, e23)) | (e21 = op2(e22, e22)) | (e21 = op2(e22, e21)) | (e21 = op2(e22, e20))), inference(cnf_transformation, [], [f6])).
fof(f3075, plain, (spl280_176 | spl280_171 | spl280_166 | spl280_161 | spl280_156), inference(avatar_split_clause, [], [f739, f2616, f2637, f2658, f2679, f2700])).
fof(f739, plain, ((e20 = op2(e23, e24)) | (e20 = op2(e23, e23)) | (e20 = op2(e23, e22)) | (e20 = op2(e23, e21)) | (e20 = op2(e23, e20))), inference(cnf_transformation, [], [f6])).
fof(f3072, plain, (spl280_237 | spl280_212 | spl280_187 | spl280_162 | spl280_137), inference(avatar_split_clause, [], [f742, f2536, f2641, f2746, f2851, f2956])).
fof(f742, plain, ((e21 = op2(e24, e23)) | (e21 = op2(e23, e23)) | (e21 = op2(e22, e23)) | (e21 = op2(e21, e23)) | (e21 = op2(e20, e23))), inference(cnf_transformation, [], [f6])).
fof(f3060, plain, (spl280_233 | spl280_208 | spl280_183 | spl280_158 | spl280_133), inference(avatar_split_clause, [], [f754, f2519, f2624, f2729, f2834, f2939])).
fof(f754, plain, ((e22 = op2(e24, e24)) | (e22 = op2(e23, e24)) | (e22 = op2(e22, e24)) | (e22 = op2(e21, e24)) | (e22 = op2(e20, e24))), inference(cnf_transformation, [], [f6])).
fof(f3055, plain, (spl280_256 | spl280_257 | spl280_258 | spl280_259 | spl280_260), inference(avatar_split_clause, [], [f708, f3052, f3048, f3044, f3040, f3036])).
fof(f708, plain, ((e24 = unit2) | (e23 = unit2) | (e22 = unit2) | (e21 = unit2) | (e20 = unit2)), inference(cnf_transformation, [], [f5])).
fof(f2866, plain, (spl280_211 | spl280_212 | spl280_213 | spl280_214 | spl280_215), inference(avatar_split_clause, [], [f681, f2863, f2859, f2855, f2851, f2847])).
fof(f681, plain, ((e24 = op2(e21, e23)) | (e23 = op2(e21, e23)) | (e22 = op2(e21, e23)) | (e21 = op2(e21, e23)) | (e20 = op2(e21, e23))), inference(cnf_transformation, [], [f4])).
fof(f4, plain, (((e24 = op2(e24, e24)) | (e23 = op2(e24, e24)) | (e22 = op2(e24, e24)) | (e21 = op2(e24, e24)) | (e20 = op2(e24, e24))) & ((e24 = op2(e24, e23)) | (e23 = op2(e24, e23)) | (e22 = op2(e24, e23)) | (e21 = op2(e24, e23)) | (e20 = op2(e24, e23))) & ((e24 = op2(e24, e22)) | (e23 = op2(e24, e22)) | (e22 = op2(e24, e22)) | (e21 = op2(e24, e22)) | (e20 = op2(e24, e22))) & ((e24 = op2(e24, e21)) | (e23 = op2(e24, e21)) | (e22 = op2(e24, e21)) | (e21 = op2(e24, e21)) | (e20 = op2(e24, e21))) & ((e24 = op2(e24, e20)) | (e23 = op2(e24, e20)) | (e22 = op2(e24, e20)) | (e21 = op2(e24, e20)) | (e20 = op2(e24, e20))) & ((e24 = op2(e23, e24)) | (e23 = op2(e23, e24)) | (e22 = op2(e23, e24)) | (e21 = op2(e23, e24)) | (e20 = op2(e23, e24))) & ((e24 = op2(e23, e23)) | (e23 = op2(e23, e23)) | (e22 = op2(e23, e23)) | (e21 = op2(e23, e23)) | (e20 = op2(e23, e23))) & ((e24 = op2(e23, e22)) | (e23 = op2(e23, e22)) | (e22 = op2(e23, e22)) | (e21 = op2(e23, e22)) | (e20 = op2(e23, e22))) & ((e24 = op2(e23, e21)) | (e23 = op2(e23, e21)) | (e22 = op2(e23, e21)) | (e21 = op2(e23, e21)) | (e20 = op2(e23, e21))) & ((e24 = op2(e23, e20)) | (e23 = op2(e23, e20)) | (e22 = op2(e23, e20)) | (e21 = op2(e23, e20)) | (e20 = op2(e23, e20))) & ((e24 = op2(e22, e24)) | (e23 = op2(e22, e24)) | (e22 = op2(e22, e24)) | (e21 = op2(e22, e24)) | (e20 = op2(e22, e24))) & ((e24 = op2(e22, e23)) | (e23 = op2(e22, e23)) | (e22 = op2(e22, e23)) | (e21 = op2(e22, e23)) | (e20 = op2(e22, e23))) & ((e24 = op2(e22, e22)) | (e23 = op2(e22, e22)) | (e22 = op2(e22, e22)) | (e21 = op2(e22, e22)) | (e20 = op2(e22, e22))) & ((e24 = op2(e22, e21)) | (e23 = op2(e22, e21)) | (e22 = op2(e22, e21)) | (e21 = op2(e22, e21)) | (e20 = op2(e22, e21))) & ((e24 = op2(e22, e20)) | (e23 = op2(e22, e20)) | (e22 = op2(e22, e20)) | (e21 = op2(e22, e20)) | (e20 = op2(e22, e20))) & ((e24 = op2(e21, e24)) | (e23 = op2(e21, e24)) | (e22 = op2(e21, e24)) | (e21 = op2(e21, e24)) | (e20 = op2(e21, e24))) & ((e24 = op2(e21, e23)) | (e23 = op2(e21, e23)) | (e22 = op2(e21, e23)) | (e21 = op2(e21, e23)) | (e20 = op2(e21, e23))) & ((e24 = op2(e21, e22)) | (e23 = op2(e21, e22)) | (e22 = op2(e21, e22)) | (e21 = op2(e21, e22)) | (e20 = op2(e21, e22))) & ((e24 = op2(e21, e21)) | (e23 = op2(e21, e21)) | (e22 = op2(e21, e21)) | (e21 = op2(e21, e21)) | (e20 = op2(e21, e21))) & ((e24 = op2(e21, e20)) | (e23 = op2(e21, e20)) | (e22 = op2(e21, e20)) | (e21 = op2(e21, e20)) | (e20 = op2(e21, e20))) & ((e24 = op2(e20, e24)) | (e23 = op2(e20, e24)) | (e22 = op2(e20, e24)) | (e21 = op2(e20, e24)) | (e20 = op2(e20, e24))) & ((e24 = op2(e20, e23)) | (e23 = op2(e20, e23)) | (e22 = op2(e20, e23)) | (e21 = op2(e20, e23)) | (e20 = op2(e20, e23))) & ((e24 = op2(e20, e22)) | (e23 = op2(e20, e22)) | (e22 = op2(e20, e22)) | (e21 = op2(e20, e22)) | (e20 = op2(e20, e22))) & ((e24 = op2(e20, e21)) | (e23 = op2(e20, e21)) | (e22 = op2(e20, e21)) | (e21 = op2(e20, e21)) | (e20 = op2(e20, e21))) & ((op2(e20, e20) = e24) | (op2(e20, e20) = e23) | (op2(e20, e20) = e22) | (op2(e20, e20) = e21) | (e20 = op2(e20, e20)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG051+1.p', ax4)).
fof(f2493, plain, (spl280_99 | spl280_94 | spl280_89 | spl280_84 | spl280_79), inference(avatar_split_clause, [], [f639, f2242, f2263, f2284, f2305, f2326])).
fof(f639, plain, ((e13 = op1(e11, e14)) | (e13 = op1(e11, e13)) | (e13 = op1(e11, e12)) | (e13 = op1(e11, e11)) | (e13 = op1(e11, e10))), inference(cnf_transformation, [], [f3])).
fof(f3, plain, (((e14 = op1(e14, e14)) | (e14 = op1(e13, e14)) | (e14 = op1(e12, e14)) | (e14 = op1(e11, e14)) | (e14 = op1(e10, e14))) & ((e14 = op1(e14, e14)) | (e14 = op1(e14, e13)) | (e14 = op1(e14, e12)) | (e14 = op1(e14, e11)) | (e14 = op1(e14, e10))) & ((e13 = op1(e14, e14)) | (e13 = op1(e13, e14)) | (e13 = op1(e12, e14)) | (e13 = op1(e11, e14)) | (e13 = op1(e10, e14))) & ((e13 = op1(e14, e14)) | (e13 = op1(e14, e13)) | (e13 = op1(e14, e12)) | (e13 = op1(e14, e11)) | (e13 = op1(e14, e10))) & ((e12 = op1(e14, e14)) | (e12 = op1(e13, e14)) | (e12 = op1(e12, e14)) | (e12 = op1(e11, e14)) | (e12 = op1(e10, e14))) & ((e12 = op1(e14, e14)) | (e12 = op1(e14, e13)) | (e12 = op1(e14, e12)) | (e12 = op1(e14, e11)) | (e12 = op1(e14, e10))) & ((e11 = op1(e14, e14)) | (e11 = op1(e13, e14)) | (e11 = op1(e12, e14)) | (e11 = op1(e11, e14)) | (e11 = op1(e10, e14))) & ((e11 = op1(e14, e14)) | (e11 = op1(e14, e13)) | (e11 = op1(e14, e12)) | (e11 = op1(e14, e11)) | (e11 = op1(e14, e10))) & ((e10 = op1(e14, e14)) | (e10 = op1(e13, e14)) | (e10 = op1(e12, e14)) | (e10 = op1(e11, e14)) | (e10 = op1(e10, e14))) & ((e10 = op1(e14, e14)) | (e10 = op1(e14, e13)) | (e10 = op1(e14, e12)) | (e10 = op1(e14, e11)) | (e10 = op1(e14, e10))) & ((e14 = op1(e14, e13)) | (e14 = op1(e13, e13)) | (e14 = op1(e12, e13)) | (e14 = op1(e11, e13)) | (e14 = op1(e10, e13))) & ((e14 = op1(e13, e14)) | (e14 = op1(e13, e13)) | (e14 = op1(e13, e12)) | (e14 = op1(e13, e11)) | (e14 = op1(e13, e10))) & ((e13 = op1(e14, e13)) | (e13 = op1(e13, e13)) | (e13 = op1(e12, e13)) | (e13 = op1(e11, e13)) | (e13 = op1(e10, e13))) & ((e13 = op1(e13, e14)) | (e13 = op1(e13, e13)) | (e13 = op1(e13, e12)) | (e13 = op1(e13, e11)) | (e13 = op1(e13, e10))) & ((e12 = op1(e14, e13)) | (e12 = op1(e13, e13)) | (e12 = op1(e12, e13)) | (e12 = op1(e11, e13)) | (e12 = op1(e10, e13))) & ((e12 = op1(e13, e14)) | (e12 = op1(e13, e13)) | (e12 = op1(e13, e12)) | (e12 = op1(e13, e11)) | (e12 = op1(e13, e10))) & ((e11 = op1(e14, e13)) | (e11 = op1(e13, e13)) | (e11 = op1(e12, e13)) | (e11 = op1(e11, e13)) | (e11 = op1(e10, e13))) & ((e11 = op1(e13, e14)) | (e11 = op1(e13, e13)) | (e11 = op1(e13, e12)) | (e11 = op1(e13, e11)) | (e11 = op1(e13, e10))) & ((e10 = op1(e14, e13)) | (e10 = op1(e13, e13)) | (e10 = op1(e12, e13)) | (e10 = op1(e11, e13)) | (e10 = op1(e10, e13))) & ((e10 = op1(e13, e14)) | (e10 = op1(e13, e13)) | (e10 = op1(e13, e12)) | (e10 = op1(e13, e11)) | (e10 = op1(e13, e10))) & ((e14 = op1(e14, e12)) | (e14 = op1(e13, e12)) | (e14 = op1(e12, e12)) | (e14 = op1(e11, e12)) | (e14 = op1(e10, e12))) & ((e14 = op1(e12, e14)) | (e14 = op1(e12, e13)) | (e14 = op1(e12, e12)) | (e14 = op1(e12, e11)) | (e14 = op1(e12, e10))) & ((e13 = op1(e14, e12)) | (e13 = op1(e13, e12)) | (e13 = op1(e12, e12)) | (e13 = op1(e11, e12)) | (e13 = op1(e10, e12))) & ((e13 = op1(e12, e14)) | (e13 = op1(e12, e13)) | (e13 = op1(e12, e12)) | (e13 = op1(e12, e11)) | (e13 = op1(e12, e10))) & ((e12 = op1(e14, e12)) | (e12 = op1(e13, e12)) | (e12 = op1(e12, e12)) | (e12 = op1(e11, e12)) | (e12 = op1(e10, e12))) & ((e12 = op1(e12, e14)) | (e12 = op1(e12, e13)) | (e12 = op1(e12, e12)) | (e12 = op1(e12, e11)) | (e12 = op1(e12, e10))) & ((e11 = op1(e14, e12)) | (e11 = op1(e13, e12)) | (e11 = op1(e12, e12)) | (e11 = op1(e11, e12)) | (e11 = op1(e10, e12))) & ((e11 = op1(e12, e14)) | (e11 = op1(e12, e13)) | (e11 = op1(e12, e12)) | (e11 = op1(e12, e11)) | (e11 = op1(e12, e10))) & ((e10 = op1(e14, e12)) | (e10 = op1(e13, e12)) | (e10 = op1(e12, e12)) | (e10 = op1(e11, e12)) | (e10 = op1(e10, e12))) & ((e10 = op1(e12, e14)) | (e10 = op1(e12, e13)) | (e10 = op1(e12, e12)) | (e10 = op1(e12, e11)) | (e10 = op1(e12, e10))) & ((e14 = op1(e14, e11)) | (e14 = op1(e13, e11)) | (e14 = op1(e12, e11)) | (e14 = op1(e11, e11)) | (e14 = op1(e10, e11))) & ((e14 = op1(e11, e14)) | (e14 = op1(e11, e13)) | (e14 = op1(e11, e12)) | (e14 = op1(e11, e11)) | (e14 = op1(e11, e10))) & ((e13 = op1(e14, e11)) | (e13 = op1(e13, e11)) | (e13 = op1(e12, e11)) | (e13 = op1(e11, e11)) | (e13 = op1(e10, e11))) & ((e13 = op1(e11, e14)) | (e13 = op1(e11, e13)) | (e13 = op1(e11, e12)) | (e13 = op1(e11, e11)) | (e13 = op1(e11, e10))) & ((e12 = op1(e14, e11)) | (e12 = op1(e13, e11)) | (e12 = op1(e12, e11)) | (e12 = op1(e11, e11)) | (e12 = op1(e10, e11))) & ((e12 = op1(e11, e14)) | (e12 = op1(e11, e13)) | (e12 = op1(e11, e12)) | (e12 = op1(e11, e11)) | (e12 = op1(e11, e10))) & ((e11 = op1(e14, e11)) | (e11 = op1(e13, e11)) | (e11 = op1(e12, e11)) | (e11 = op1(e11, e11)) | (e11 = op1(e10, e11))) & ((e11 = op1(e11, e14)) | (e11 = op1(e11, e13)) | (e11 = op1(e11, e12)) | (e11 = op1(e11, e11)) | (e11 = op1(e11, e10))) & ((e10 = op1(e14, e11)) | (e10 = op1(e13, e11)) | (e10 = op1(e12, e11)) | (e10 = op1(e11, e11)) | (e10 = op1(e10, e11))) & ((e10 = op1(e11, e14)) | (e10 = op1(e11, e13)) | (e10 = op1(e11, e12)) | (e10 = op1(e11, e11)) | (e10 = op1(e11, e10))) & ((e14 = op1(e14, e10)) | (e14 = op1(e13, e10)) | (e14 = op1(e12, e10)) | (e14 = op1(e11, e10)) | (op1(e10, e10) = e14)) & ((e14 = op1(e10, e14)) | (e14 = op1(e10, e13)) | (e14 = op1(e10, e12)) | (e14 = op1(e10, e11)) | (op1(e10, e10) = e14)) & ((e13 = op1(e14, e10)) | (e13 = op1(e13, e10)) | (e13 = op1(e12, e10)) | (e13 = op1(e11, e10)) | (op1(e10, e10) = e13)) & ((e13 = op1(e10, e14)) | (e13 = op1(e10, e13)) | (e13 = op1(e10, e12)) | (e13 = op1(e10, e11)) | (op1(e10, e10) = e13)) & ((e12 = op1(e14, e10)) | (e12 = op1(e13, e10)) | (e12 = op1(e12, e10)) | (e12 = op1(e11, e10)) | (op1(e10, e10) = e12)) & ((e12 = op1(e10, e14)) | (e12 = op1(e10, e13)) | (e12 = op1(e10, e12)) | (e12 = op1(e10, e11)) | (op1(e10, e10) = e12)) & ((e11 = op1(e14, e10)) | (e11 = op1(e13, e10)) | (e11 = op1(e12, e10)) | (e11 = op1(e11, e10)) | (op1(e10, e10) = e11)) & ((e11 = op1(e10, e14)) | (e11 = op1(e10, e13)) | (e11 = op1(e10, e12)) | (e11 = op1(e10, e11)) | (op1(e10, e10) = e11)) & ((e10 = op1(e14, e10)) | (e10 = op1(e13, e10)) | (e10 = op1(e12, e10)) | (e10 = op1(e11, e10)) | (e10 = op1(e10, e10))) & ((e10 = op1(e10, e14)) | (e10 = op1(e10, e13)) | (e10 = op1(e10, e12)) | (e10 = op1(e10, e11)) | (e10 = op1(e10, e10)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG051+1.p', ax3)).
fof(f2491, plain, (spl280_100 | spl280_95 | spl280_90 | spl280_85 | spl280_80), inference(avatar_split_clause, [], [f641, f2246, f2267, f2288, f2309, f2330])).
fof(f641, plain, ((e14 = op1(e11, e14)) | (e14 = op1(e11, e13)) | (e14 = op1(e11, e12)) | (e14 = op1(e11, e11)) | (e14 = op1(e11, e10))), inference(cnf_transformation, [], [f3])).
fof(f2490, plain, (spl280_120 | spl280_95 | spl280_70 | spl280_45 | spl280_20), inference(avatar_split_clause, [], [f642, f1994, f2099, f2204, f2309, f2414])).
fof(f642, plain, ((e14 = op1(e14, e11)) | (e14 = op1(e13, e11)) | (e14 = op1(e12, e11)) | (e14 = op1(e11, e11)) | (e14 = op1(e10, e11))), inference(cnf_transformation, [], [f3])).
fof(f2488, plain, (spl280_111 | spl280_86 | spl280_61 | spl280_36 | spl280_11), inference(avatar_split_clause, [], [f644, f1957, f2062, f2167, f2272, f2377])).
fof(f644, plain, ((e10 = op1(e14, e12)) | (e10 = op1(e13, e12)) | (e10 = op1(e12, e12)) | (e10 = op1(e11, e12)) | (e10 = op1(e10, e12))), inference(cnf_transformation, [], [f3])).
fof(f2479, plain, (spl280_46 | spl280_41 | spl280_36 | spl280_31 | spl280_26), inference(avatar_split_clause, [], [f653, f2020, f2041, f2062, f2083, f2104])).
fof(f653, plain, ((e10 = op1(e13, e14)) | (e10 = op1(e13, e13)) | (e10 = op1(e13, e12)) | (e10 = op1(e13, e11)) | (e10 = op1(e13, e10))), inference(cnf_transformation, [], [f3])).
fof(f2478, plain, (spl280_106 | spl280_81 | spl280_56 | spl280_31 | spl280_6), inference(avatar_split_clause, [], [f654, f1936, f2041, f2146, f2251, f2356])).
fof(f654, plain, ((e10 = op1(e14, e13)) | (e10 = op1(e13, e13)) | (e10 = op1(e12, e13)) | (e10 = op1(e11, e13)) | (e10 = op1(e10, e13))), inference(cnf_transformation, [], [f3])).
fof(f2467, plain, (spl280_22 | spl280_17 | spl280_12 | spl280_7 | spl280_2), inference(avatar_split_clause, [], [f665, f1919, f1940, f1961, f1982, f2003])).
fof(f665, plain, ((e11 = op1(e14, e14)) | (e11 = op1(e14, e13)) | (e11 = op1(e14, e12)) | (e11 = op1(e14, e11)) | (e11 = op1(e14, e10))), inference(cnf_transformation, [], [f3])).
fof(f2466, plain, (spl280_102 | spl280_77 | spl280_52 | spl280_27 | spl280_2), inference(avatar_split_clause, [], [f666, f1919, f2024, f2129, f2234, f2339])).
fof(f666, plain, ((e11 = op1(e14, e14)) | (e11 = op1(e13, e14)) | (e11 = op1(e12, e14)) | (e11 = op1(e11, e14)) | (e11 = op1(e10, e14))), inference(cnf_transformation, [], [f3])).
fof(f2464, plain, (spl280_103 | spl280_78 | spl280_53 | spl280_28 | spl280_3), inference(avatar_split_clause, [], [f668, f1923, f2028, f2133, f2238, f2343])).
fof(f668, plain, ((e12 = op1(e14, e14)) | (e12 = op1(e13, e14)) | (e12 = op1(e12, e14)) | (e12 = op1(e11, e14)) | (e12 = op1(e10, e14))), inference(cnf_transformation, [], [f3])).
fof(f2463, plain, (spl280_24 | spl280_19 | spl280_14 | spl280_9 | spl280_4), inference(avatar_split_clause, [], [f669, f1927, f1948, f1969, f1990, f2011])).
fof(f669, plain, ((e13 = op1(e14, e14)) | (e13 = op1(e14, e13)) | (e13 = op1(e14, e12)) | (e13 = op1(e14, e11)) | (e13 = op1(e14, e10))), inference(cnf_transformation, [], [f3])).
fof(f2459, plain, (spl280_126 | spl280_127 | spl280_128 | spl280_129 | spl280_130), inference(avatar_split_clause, [], [f622, f2456, f2452, f2448, f2444, f2440])).
fof(f622, plain, ((e14 = unit1) | (e13 = unit1) | (e12 = unit1) | (e11 = unit1) | (e10 = unit1)), inference(cnf_transformation, [], [f2])).
fof(f2270, plain, (spl280_81 | spl280_82 | spl280_83 | spl280_84 | spl280_85), inference(avatar_split_clause, [], [f595, f2267, f2263, f2259, f2255, f2251])).
fof(f595, plain, ((e14 = op1(e11, e13)) | (e13 = op1(e11, e13)) | (e12 = op1(e11, e13)) | (e11 = op1(e11, e13)) | (e10 = op1(e11, e13))), inference(cnf_transformation, [], [f1])).
fof(f1, plain, (((e14 = op1(e14, e14)) | (e13 = op1(e14, e14)) | (e12 = op1(e14, e14)) | (e11 = op1(e14, e14)) | (e10 = op1(e14, e14))) & ((e14 = op1(e14, e13)) | (e13 = op1(e14, e13)) | (e12 = op1(e14, e13)) | (e11 = op1(e14, e13)) | (e10 = op1(e14, e13))) & ((e14 = op1(e14, e12)) | (e13 = op1(e14, e12)) | (e12 = op1(e14, e12)) | (e11 = op1(e14, e12)) | (e10 = op1(e14, e12))) & ((e14 = op1(e14, e11)) | (e13 = op1(e14, e11)) | (e12 = op1(e14, e11)) | (e11 = op1(e14, e11)) | (e10 = op1(e14, e11))) & ((e14 = op1(e14, e10)) | (e13 = op1(e14, e10)) | (e12 = op1(e14, e10)) | (e11 = op1(e14, e10)) | (e10 = op1(e14, e10))) & ((e14 = op1(e13, e14)) | (e13 = op1(e13, e14)) | (e12 = op1(e13, e14)) | (e11 = op1(e13, e14)) | (e10 = op1(e13, e14))) & ((e14 = op1(e13, e13)) | (e13 = op1(e13, e13)) | (e12 = op1(e13, e13)) | (e11 = op1(e13, e13)) | (e10 = op1(e13, e13))) & ((e14 = op1(e13, e12)) | (e13 = op1(e13, e12)) | (e12 = op1(e13, e12)) | (e11 = op1(e13, e12)) | (e10 = op1(e13, e12))) & ((e14 = op1(e13, e11)) | (e13 = op1(e13, e11)) | (e12 = op1(e13, e11)) | (e11 = op1(e13, e11)) | (e10 = op1(e13, e11))) & ((e14 = op1(e13, e10)) | (e13 = op1(e13, e10)) | (e12 = op1(e13, e10)) | (e11 = op1(e13, e10)) | (e10 = op1(e13, e10))) & ((e14 = op1(e12, e14)) | (e13 = op1(e12, e14)) | (e12 = op1(e12, e14)) | (e11 = op1(e12, e14)) | (e10 = op1(e12, e14))) & ((e14 = op1(e12, e13)) | (e13 = op1(e12, e13)) | (e12 = op1(e12, e13)) | (e11 = op1(e12, e13)) | (e10 = op1(e12, e13))) & ((e14 = op1(e12, e12)) | (e13 = op1(e12, e12)) | (e12 = op1(e12, e12)) | (e11 = op1(e12, e12)) | (e10 = op1(e12, e12))) & ((e14 = op1(e12, e11)) | (e13 = op1(e12, e11)) | (e12 = op1(e12, e11)) | (e11 = op1(e12, e11)) | (e10 = op1(e12, e11))) & ((e14 = op1(e12, e10)) | (e13 = op1(e12, e10)) | (e12 = op1(e12, e10)) | (e11 = op1(e12, e10)) | (e10 = op1(e12, e10))) & ((e14 = op1(e11, e14)) | (e13 = op1(e11, e14)) | (e12 = op1(e11, e14)) | (e11 = op1(e11, e14)) | (e10 = op1(e11, e14))) & ((e14 = op1(e11, e13)) | (e13 = op1(e11, e13)) | (e12 = op1(e11, e13)) | (e11 = op1(e11, e13)) | (e10 = op1(e11, e13))) & ((e14 = op1(e11, e12)) | (e13 = op1(e11, e12)) | (e12 = op1(e11, e12)) | (e11 = op1(e11, e12)) | (e10 = op1(e11, e12))) & ((e14 = op1(e11, e11)) | (e13 = op1(e11, e11)) | (e12 = op1(e11, e11)) | (e11 = op1(e11, e11)) | (e10 = op1(e11, e11))) & ((e14 = op1(e11, e10)) | (e13 = op1(e11, e10)) | (e12 = op1(e11, e10)) | (e11 = op1(e11, e10)) | (e10 = op1(e11, e10))) & ((e14 = op1(e10, e14)) | (e13 = op1(e10, e14)) | (e12 = op1(e10, e14)) | (e11 = op1(e10, e14)) | (e10 = op1(e10, e14))) & ((e14 = op1(e10, e13)) | (e13 = op1(e10, e13)) | (e12 = op1(e10, e13)) | (e11 = op1(e10, e13)) | (e10 = op1(e10, e13))) & ((e14 = op1(e10, e12)) | (e13 = op1(e10, e12)) | (e12 = op1(e10, e12)) | (e11 = op1(e10, e12)) | (e10 = op1(e10, e12))) & ((e14 = op1(e10, e11)) | (e13 = op1(e10, e11)) | (e12 = op1(e10, e11)) | (e11 = op1(e10, e11)) | (e10 = op1(e10, e11))) & ((op1(e10, e10) = e14) | (op1(e10, e10) = e13) | (op1(e10, e10) = e12) | (op1(e10, e10) = e11) | (e10 = op1(e10, e10)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG051+1.p', ax1)).