fof(f9779, plain, $false, inference(avatar_sat_refutation, [], [f1074, f1176, f1193, f1431, f1465, f1482, f1528, f1536, f1545, f1546, f1547, f1548, f1557, f1565, f1573, f1574, f1575, f1576, f1580, f1581, f1582, f1583, f1779, f1802, f2157, f3866, f3867, f4008, f4010, f4017, f4023, f4088, f4127, f4286, f4550, f4631, f4673, f4765, f4823, f5082, f5097, f5127, f5136, f5202, f5309, f5370, f5467, f5604, f5663, f5735, f6015, f6046, f6051, f6234, f6263, f6402, f6505, f6587, f6592, f6644, f6674, f6679, f6718, f6721, f6806, f6983, f7057, f7058, f7205, f7395, f8002, f8299, f8301, f8302, f8303, f8380, f8411, f8487, f8549, f8606, f8662, f8722, f8784, f8838, f8889, f8951, f9013, f9077, f9150, f9480, f9778])).
fof(f9778, plain, (~ spl42_1 | ~ spl42_2), inference(avatar_contradiction_clause, [], [f9777])).
fof(f9777, plain, ($false | (~ spl42_1 | ~ spl42_2)), inference(subsumption_resolution, [], [f9776, f482])).
fof(f482, plain, ~ (e10 = e11), inference(cnf_transformation, [], [f19])).
fof(f19, plain, (~ (e12 = e13) & ~ (e11 = e13) & ~ (e11 = e12) & ~ (e10 = e13) & ~ (e10 = e12) & ~ (e10 = e11)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG015+1.p', ax19)).
fof(f9776, plain, ((e10 = e11) | (~ spl42_1 | ~ spl42_2)), inference(forward_demodulation, [], [f793, f789])).
fof(f789, plain, ((e10 = op1(e13, e13)) | ~ spl42_1), inference(avatar_component_clause, [], [f787])).
fof(f787, plain, (spl42_1 <=> (e10 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_1])])).
fof(f793, plain, ((e11 = op1(e13, e13)) | ~ spl42_2), inference(avatar_component_clause, [], [f791])).
fof(f791, plain, (spl42_2 <=> (e11 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_2])])).
fof(f9480, plain, (~ spl42_1 | ~ spl42_18 | ~ spl42_21 | ~ spl42_100 | ~ spl42_112 | ~ spl42_180 | ~ spl42_194 | ~ spl42_229 | spl42_233), inference(avatar_contradiction_clause, [], [f9479])).
fof(f9479, plain, ($false | (~ spl42_1 | ~ spl42_18 | ~ spl42_21 | ~ spl42_100 | ~ spl42_112 | ~ spl42_180 | ~ spl42_194 | ~ spl42_229 | spl42_233)), inference(subsumption_resolution, [], [f9478, f1260])).
fof(f1260, plain, ((e23 = op2(e22, e21)) | ~ spl42_112), inference(avatar_component_clause, [], [f1258])).
fof(f1258, plain, (spl42_112 <=> (e23 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_112])])).
fof(f9478, plain, (~ (e23 = op2(e22, e21)) | (~ spl42_1 | ~ spl42_18 | ~ spl42_21 | ~ spl42_100 | ~ spl42_180 | ~ spl42_194 | ~ spl42_229 | spl42_233)), inference(forward_demodulation, [], [f9477, f9363])).
fof(f9363, plain, ((e23 = h12(op1(e10, e12))) | (~ spl42_1 | ~ spl42_18 | ~ spl42_21 | ~ spl42_100 | ~ spl42_180 | ~ spl42_229)), inference(backward_demodulation, [], [f9183, f7588])).
fof(f7588, plain, ((op1(e10, e12) = op1(e11, e13)) | (~ spl42_1 | ~ spl42_18 | ~ spl42_21)), inference(backward_demodulation, [], [f5835, f5860])).
fof(f5860, plain, ((op1(e10, e12) = op1(e12, e10)) | ~ spl42_21), inference(forward_demodulation, [], [f188, f874])).
fof(f874, plain, ((e10 = op1(e12, e12)) | ~ spl42_21), inference(avatar_component_clause, [], [f872])).
fof(f872, plain, (spl42_21 <=> (e10 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_21])])).
fof(f188, plain, (op1(op1(e12, e12), e12) = op1(e12, op1(e12, e12))), inference(cnf_transformation, [], [f2])).
fof(f2, plain, ((op1(op1(e13, e13), e13) = op1(e13, op1(e13, e13))) & (op1(op1(e13, e13), e12) = op1(e13, op1(e13, e12))) & (op1(op1(e13, e13), e11) = op1(e13, op1(e13, e11))) & (op1(op1(e13, e13), e10) = op1(e13, op1(e13, e10))) & (op1(op1(e13, e12), e13) = op1(e13, op1(e12, e13))) & (op1(op1(e13, e12), e12) = op1(e13, op1(e12, e12))) & (op1(op1(e13, e12), e11) = op1(e13, op1(e12, e11))) & (op1(op1(e13, e12), e10) = op1(e13, op1(e12, e10))) & (op1(op1(e13, e11), e13) = op1(e13, op1(e11, e13))) & (op1(op1(e13, e11), e12) = op1(e13, op1(e11, e12))) & (op1(op1(e13, e11), e11) = op1(e13, op1(e11, e11))) & (op1(op1(e13, e11), e10) = op1(e13, op1(e11, e10))) & (op1(op1(e13, e10), e13) = op1(e13, op1(e10, e13))) & (op1(op1(e13, e10), e12) = op1(e13, op1(e10, e12))) & (op1(op1(e13, e10), e11) = op1(e13, op1(e10, e11))) & (op1(op1(e13, e10), e10) = op1(e13, op1(e10, e10))) & (op1(op1(e12, e13), e13) = op1(e12, op1(e13, e13))) & (op1(op1(e12, e13), e12) = op1(e12, op1(e13, e12))) & (op1(op1(e12, e13), e11) = op1(e12, op1(e13, e11))) & (op1(op1(e12, e13), e10) = op1(e12, op1(e13, e10))) & (op1(op1(e12, e12), e13) = op1(e12, op1(e12, e13))) & (op1(op1(e12, e12), e12) = op1(e12, op1(e12, e12))) & (op1(op1(e12, e12), e11) = op1(e12, op1(e12, e11))) & (op1(op1(e12, e12), e10) = op1(e12, op1(e12, e10))) & (op1(op1(e12, e11), e13) = op1(e12, op1(e11, e13))) & (op1(op1(e12, e11), e12) = op1(e12, op1(e11, e12))) & (op1(op1(e12, e11), e11) = op1(e12, op1(e11, e11))) & (op1(op1(e12, e11), e10) = op1(e12, op1(e11, e10))) & (op1(op1(e12, e10), e13) = op1(e12, op1(e10, e13))) & (op1(op1(e12, e10), e12) = op1(e12, op1(e10, e12))) & (op1(op1(e12, e10), e11) = op1(e12, op1(e10, e11))) & (op1(op1(e12, e10), e10) = op1(e12, op1(e10, e10))) & (op1(op1(e11, e13), e13) = op1(e11, op1(e13, e13))) & (op1(op1(e11, e13), e12) = op1(e11, op1(e13, e12))) & (op1(op1(e11, e13), e11) = op1(e11, op1(e13, e11))) & (op1(op1(e11, e13), e10) = op1(e11, op1(e13, e10))) & (op1(op1(e11, e12), e13) = op1(e11, op1(e12, e13))) & (op1(op1(e11, e12), e12) = op1(e11, op1(e12, e12))) & (op1(op1(e11, e12), e11) = op1(e11, op1(e12, e11))) & (op1(op1(e11, e12), e10) = op1(e11, op1(e12, e10))) & (op1(op1(e11, e11), e13) = op1(e11, op1(e11, e13))) & (op1(op1(e11, e11), e12) = op1(e11, op1(e11, e12))) & (op1(op1(e11, e11), e11) = op1(e11, op1(e11, e11))) & (op1(op1(e11, e11), e10) = op1(e11, op1(e11, e10))) & (op1(op1(e11, e10), e13) = op1(e11, op1(e10, e13))) & (op1(op1(e11, e10), e12) = op1(e11, op1(e10, e12))) & (op1(op1(e11, e10), e11) = op1(e11, op1(e10, e11))) & (op1(op1(e11, e10), e10) = op1(e11, op1(e10, e10))) & (op1(op1(e10, e13), e13) = op1(e10, op1(e13, e13))) & (op1(op1(e10, e13), e12) = op1(e10, op1(e13, e12))) & (op1(op1(e10, e13), e11) = op1(e10, op1(e13, e11))) & (op1(op1(e10, e13), e10) = op1(e10, op1(e13, e10))) & (op1(op1(e10, e12), e13) = op1(e10, op1(e12, e13))) & (op1(op1(e10, e12), e12) = op1(e10, op1(e12, e12))) & (op1(op1(e10, e12), e11) = op1(e10, op1(e12, e11))) & (op1(op1(e10, e12), e10) = op1(e10, op1(e12, e10))) & (op1(op1(e10, e11), e13) = op1(e10, op1(e11, e13))) & (op1(op1(e10, e11), e12) = op1(e10, op1(e11, e12))) & (op1(op1(e10, e11), e11) = op1(e10, op1(e11, e11))) & (op1(op1(e10, e11), e10) = op1(e10, op1(e11, e10))) & (op1(op1(e10, e10), e13) = op1(e10, op1(e10, e13))) & (op1(op1(e10, e10), e12) = op1(e10, op1(e10, e12))) & (op1(op1(e10, e10), e11) = op1(e10, op1(e10, e11))) & (op1(op1(e10, e10), e10) = op1(e10, op1(e10, e10)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG015+1.p', ax2)).
fof(f5835, plain, ((op1(e11, e13) = op1(e12, e10)) | (~ spl42_1 | ~ spl42_18)), inference(forward_demodulation, [], [f5834, f861])).
fof(f861, plain, ((e11 = op1(e12, e13)) | ~ spl42_18), inference(avatar_component_clause, [], [f859])).
fof(f859, plain, (spl42_18 <=> (e11 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_18])])).
fof(f5834, plain, ((op1(e12, e10) = op1(op1(e12, e13), e13)) | ~ spl42_1), inference(forward_demodulation, [], [f193, f789])).
fof(f193, plain, (op1(op1(e12, e13), e13) = op1(e12, op1(e13, e13))), inference(cnf_transformation, [], [f2])).
fof(f9183, plain, ((e23 = h12(op1(e11, e13))) | (~ spl42_100 | ~ spl42_180 | ~ spl42_229)), inference(forward_demodulation, [], [f8839, f1209])).
fof(f1209, plain, ((e23 = op2(e23, e20)) | ~ spl42_100), inference(avatar_component_clause, [], [f1207])).
fof(f1207, plain, (spl42_100 <=> (e23 = op2(e23, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_100])])).
fof(f8839, plain, ((op2(e23, e20) = h12(op1(e11, e13))) | (~ spl42_180 | ~ spl42_229)), inference(forward_demodulation, [], [f2131, f1800])).
fof(f1800, plain, ((e20 = h2(e10)) | ~ spl42_180), inference(avatar_component_clause, [], [f1799])).
fof(f1799, plain, (spl42_180 <=> (e20 = h2(e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_180])])).
fof(f2131, plain, ((h12(op1(e11, e13)) = op2(e23, h2(e10))) | ~ spl42_229), inference(avatar_component_clause, [], [f2130])).
fof(f2130, plain, (spl42_229 <=> (h12(op1(e11, e13)) = op2(e23, h2(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl42_229])])).
fof(f9477, plain, (~ (op2(e22, e21) = h12(op1(e10, e12))) | (~ spl42_194 | spl42_233)), inference(forward_demodulation, [], [f2148, f1883])).
fof(f1883, plain, ((e21 = h9(e11)) | ~ spl42_194), inference(avatar_component_clause, [], [f1882])).
fof(f1882, plain, (spl42_194 <=> (e21 = h9(e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_194])])).
fof(f2148, plain, (~ (h12(op1(e10, e12)) = op2(e22, h9(e11))) | spl42_233), inference(avatar_component_clause, [], [f2146])).
fof(f2146, plain, (spl42_233 <=> (h12(op1(e10, e12)) = op2(e22, h9(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl42_233])])).
fof(f9150, plain, (~ spl42_31 | ~ spl42_100 | ~ spl42_180 | spl42_228), inference(avatar_contradiction_clause, [], [f9149])).
fof(f9149, plain, ($false | (~ spl42_31 | ~ spl42_100 | ~ spl42_180 | spl42_228)), inference(subsumption_resolution, [], [f9148, f1209])).
fof(f9148, plain, (~ (e23 = op2(e23, e20)) | (~ spl42_31 | ~ spl42_180 | spl42_228)), inference(forward_demodulation, [], [f9147, f590])).
fof(f590, plain, (e23 = h12(e12)), inference(cnf_transformation, [], [f37])).
fof(f37, plain, ((op2(e23, e22) = h12(e11)) & (op2(e22, e22) = h12(e10)) & (e22 = h12(e13)) & (e23 = h12(e12))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG015+1.p', ax37)).
fof(f9147, plain, (~ (op2(e23, e20) = h12(e12)) | (~ spl42_31 | ~ spl42_180 | spl42_228)), inference(forward_demodulation, [], [f9146, f916])).
fof(f916, plain, ((e12 = op1(e12, e10)) | ~ spl42_31), inference(avatar_component_clause, [], [f914])).
fof(f914, plain, (spl42_31 <=> (e12 = op1(e12, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_31])])).
fof(f9146, plain, (~ (op2(e23, e20) = h12(op1(e12, e10))) | (~ spl42_180 | spl42_228)), inference(forward_demodulation, [], [f2128, f1800])).
fof(f2128, plain, (~ (h12(op1(e12, e10)) = op2(e23, h2(e10))) | spl42_228), inference(avatar_component_clause, [], [f2126])).
fof(f2126, plain, (spl42_228 <=> (h12(op1(e12, e10)) = op2(e23, h2(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl42_228])])).
fof(f9077, plain, (~ spl42_61 | ~ spl42_145 | ~ spl42_180 | spl42_235), inference(avatar_contradiction_clause, [], [f9076])).
fof(f9076, plain, ($false | (~ spl42_61 | ~ spl42_145 | ~ spl42_180 | spl42_235)), inference(subsumption_resolution, [], [f9075, f1401])).
fof(f1401, plain, ((e20 = op2(e20, e20)) | ~ spl42_145), inference(avatar_component_clause, [], [f1399])).
fof(f1399, plain, (spl42_145 <=> (e20 = op2(e20, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_145])])).
fof(f9075, plain, (~ (e20 = op2(e20, e20)) | (~ spl42_61 | ~ spl42_180 | spl42_235)), inference(forward_demodulation, [], [f9074, f6821])).
fof(f6821, plain, ((e20 = h12(e10)) | ~ spl42_180), inference(backward_demodulation, [], [f1767, f1800])).
fof(f1767, plain, (h2(e10) = h12(e10)), inference(forward_demodulation, [], [f592, f552])).
fof(f552, plain, (op2(e22, e22) = h2(e10)), inference(cnf_transformation, [], [f27])).
fof(f27, plain, ((op2(e20, e22) = h2(e11)) & (op2(e22, e22) = h2(e10)) & (e22 = h2(e13)) & (e20 = h2(e12))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG015+1.p', ax27)).
fof(f592, plain, (op2(e22, e22) = h12(e10)), inference(cnf_transformation, [], [f37])).
fof(f9074, plain, (~ (op2(e20, e20) = h12(e10)) | (~ spl42_61 | ~ spl42_180 | spl42_235)), inference(forward_demodulation, [], [f9073, f1800])).
fof(f9073, plain, (~ (h12(e10) = op2(h2(e10), h2(e10))) | (~ spl42_61 | spl42_235)), inference(forward_demodulation, [], [f2156, f1044])).
fof(f1044, plain, ((e10 = op1(e10, e10)) | ~ spl42_61), inference(avatar_component_clause, [], [f1042])).
fof(f1042, plain, (spl42_61 <=> (e10 = op1(e10, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_61])])).
fof(f2156, plain, (~ (op2(h2(e10), h2(e10)) = h12(op1(e10, e10))) | spl42_235), inference(avatar_component_clause, [], [f2154])).
fof(f2154, plain, (spl42_235 <=> (op2(h2(e10), h2(e10)) = h12(op1(e10, e10)))), introduced(avatar_definition, [new_symbols(naming, [spl42_235])])).
fof(f9013, plain, (~ spl42_58 | ~ spl42_142 | ~ spl42_176 | ~ spl42_180 | spl42_234), inference(avatar_contradiction_clause, [], [f9012])).
fof(f9012, plain, ($false | (~ spl42_58 | ~ spl42_142 | ~ spl42_176 | ~ spl42_180 | spl42_234)), inference(subsumption_resolution, [], [f9011, f1388])).
fof(f1388, plain, ((e21 = op2(e20, e21)) | ~ spl42_142), inference(avatar_component_clause, [], [f1386])).
fof(f1386, plain, (spl42_142 <=> (e21 = op2(e20, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_142])])).
fof(f9011, plain, (~ (e21 = op2(e20, e21)) | (~ spl42_58 | ~ spl42_176 | ~ spl42_180 | spl42_234)), inference(forward_demodulation, [], [f9010, f1777])).
fof(f1777, plain, ((e21 = h12(e11)) | ~ spl42_176), inference(avatar_component_clause, [], [f1776])).
fof(f1776, plain, (spl42_176 <=> (e21 = h12(e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_176])])).
fof(f9010, plain, (~ (op2(e20, e21) = h12(e11)) | (~ spl42_58 | ~ spl42_176 | ~ spl42_180 | spl42_234)), inference(forward_demodulation, [], [f9009, f1031])).
fof(f1031, plain, ((e11 = op1(e10, e11)) | ~ spl42_58), inference(avatar_component_clause, [], [f1029])).
fof(f1029, plain, (spl42_58 <=> (e11 = op1(e10, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_58])])).
fof(f9009, plain, (~ (op2(e20, e21) = h12(op1(e10, e11))) | (~ spl42_176 | ~ spl42_180 | spl42_234)), inference(forward_demodulation, [], [f9008, f1800])).
fof(f9008, plain, (~ (h12(op1(e10, e11)) = op2(h2(e10), e21)) | (~ spl42_176 | spl42_234)), inference(forward_demodulation, [], [f2152, f1777])).
fof(f2152, plain, (~ (h12(op1(e10, e11)) = op2(h2(e10), h12(e11))) | spl42_234), inference(avatar_component_clause, [], [f2150])).
fof(f2150, plain, (spl42_234 <=> (h12(op1(e10, e11)) = op2(h2(e10), h12(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl42_234])])).
fof(f8951, plain, (~ spl42_21 | ~ spl42_180 | ~ spl42_198 | spl42_226), inference(avatar_contradiction_clause, [], [f8950])).
fof(f8950, plain, ($false | (~ spl42_21 | ~ spl42_180 | ~ spl42_198 | spl42_226)), inference(subsumption_resolution, [], [f8949, f1906])).
fof(f1906, plain, ((e20 = h3(e10)) | ~ spl42_198), inference(avatar_component_clause, [], [f1905])).
fof(f1905, plain, (spl42_198 <=> (e20 = h3(e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_198])])).
fof(f8949, plain, (~ (e20 = h3(e10)) | (~ spl42_21 | ~ spl42_180 | spl42_226)), inference(forward_demodulation, [], [f8948, f6821])).
fof(f8948, plain, (~ (h3(e10) = h12(e10)) | (~ spl42_21 | spl42_226)), inference(forward_demodulation, [], [f2120, f874])).
fof(f2120, plain, (~ (h3(e10) = h12(op1(e12, e12))) | spl42_226), inference(avatar_component_clause, [], [f2118])).
fof(f2118, plain, (spl42_226 <=> (h3(e10) = h12(op1(e12, e12)))), introduced(avatar_definition, [new_symbols(naming, [spl42_226])])).
fof(f8889, plain, (~ spl42_46 | ~ spl42_130 | ~ spl42_176 | ~ spl42_180 | spl42_231), inference(avatar_contradiction_clause, [], [f8888])).
fof(f8888, plain, ($false | (~ spl42_46 | ~ spl42_130 | ~ spl42_176 | ~ spl42_180 | spl42_231)), inference(subsumption_resolution, [], [f8887, f1337])).
fof(f1337, plain, ((e21 = op2(e21, e20)) | ~ spl42_130), inference(avatar_component_clause, [], [f1335])).
fof(f1335, plain, (spl42_130 <=> (e21 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_130])])).
fof(f8887, plain, (~ (e21 = op2(e21, e20)) | (~ spl42_46 | ~ spl42_176 | ~ spl42_180 | spl42_231)), inference(forward_demodulation, [], [f8886, f1777])).
fof(f8886, plain, (~ (op2(e21, e20) = h12(e11)) | (~ spl42_46 | ~ spl42_176 | ~ spl42_180 | spl42_231)), inference(forward_demodulation, [], [f8885, f980])).
fof(f980, plain, ((e11 = op1(e11, e10)) | ~ spl42_46), inference(avatar_component_clause, [], [f978])).
fof(f978, plain, (spl42_46 <=> (e11 = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_46])])).
fof(f8885, plain, (~ (op2(e21, e20) = h12(op1(e11, e10))) | (~ spl42_176 | ~ spl42_180 | spl42_231)), inference(forward_demodulation, [], [f8884, f1777])).
fof(f8884, plain, (~ (h12(op1(e11, e10)) = op2(h12(e11), e20)) | (~ spl42_180 | spl42_231)), inference(forward_demodulation, [], [f2140, f1800])).
fof(f2140, plain, (~ (h12(op1(e11, e10)) = op2(h12(e11), h2(e10))) | spl42_231), inference(avatar_component_clause, [], [f2138])).
fof(f2138, plain, (spl42_231 <=> (h12(op1(e11, e10)) = op2(h12(e11), h2(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl42_231])])).
fof(f8838, plain, (~ spl42_40 | ~ spl42_95 | ~ spl42_194 | spl42_230), inference(avatar_contradiction_clause, [], [f8837])).
fof(f8837, plain, ($false | (~ spl42_40 | ~ spl42_95 | ~ spl42_194 | spl42_230)), inference(subsumption_resolution, [], [f8836, f1188])).
fof(f1188, plain, ((e22 = op2(e23, e21)) | ~ spl42_95), inference(avatar_component_clause, [], [f1186])).
fof(f1186, plain, (spl42_95 <=> (e22 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_95])])).
fof(f8836, plain, (~ (e22 = op2(e23, e21)) | (~ spl42_40 | ~ spl42_194 | spl42_230)), inference(forward_demodulation, [], [f8835, f591])).
fof(f591, plain, (e22 = h12(e13)), inference(cnf_transformation, [], [f37])).
fof(f8835, plain, (~ (op2(e23, e21) = h12(e13)) | (~ spl42_40 | ~ spl42_194 | spl42_230)), inference(forward_demodulation, [], [f8834, f954])).
fof(f954, plain, ((e13 = op1(e11, e12)) | ~ spl42_40), inference(avatar_component_clause, [], [f952])).
fof(f952, plain, (spl42_40 <=> (e13 = op1(e11, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_40])])).
fof(f8834, plain, (~ (op2(e23, e21) = h12(op1(e11, e12))) | (~ spl42_194 | spl42_230)), inference(forward_demodulation, [], [f2136, f1883])).
fof(f2136, plain, (~ (h12(op1(e11, e12)) = op2(e23, h9(e11))) | spl42_230), inference(avatar_component_clause, [], [f2134])).
fof(f2134, plain, (spl42_230 <=> (h12(op1(e11, e12)) = op2(e23, h9(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl42_230])])).
fof(f8784, plain, (~ spl42_35 | ~ spl42_100 | ~ spl42_180 | spl42_229), inference(avatar_contradiction_clause, [], [f8783])).
fof(f8783, plain, ($false | (~ spl42_35 | ~ spl42_100 | ~ spl42_180 | spl42_229)), inference(subsumption_resolution, [], [f8782, f1209])).
fof(f8782, plain, (~ (e23 = op2(e23, e20)) | (~ spl42_35 | ~ spl42_180 | spl42_229)), inference(forward_demodulation, [], [f8781, f590])).
fof(f8781, plain, (~ (op2(e23, e20) = h12(e12)) | (~ spl42_35 | ~ spl42_180 | spl42_229)), inference(forward_demodulation, [], [f8780, f933])).
fof(f933, plain, ((e12 = op1(e11, e13)) | ~ spl42_35), inference(avatar_component_clause, [], [f931])).
fof(f931, plain, (spl42_35 <=> (e12 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_35])])).
fof(f8780, plain, (~ (op2(e23, e20) = h12(op1(e11, e13))) | (~ spl42_180 | spl42_229)), inference(forward_demodulation, [], [f2132, f1800])).
fof(f2132, plain, (~ (h12(op1(e11, e13)) = op2(e23, h2(e10))) | spl42_229), inference(avatar_component_clause, [], [f2130])).
fof(f8722, plain, (~ spl42_11 | ~ spl42_112 | ~ spl42_176 | spl42_223), inference(avatar_contradiction_clause, [], [f8721])).
fof(f8721, plain, ($false | (~ spl42_11 | ~ spl42_112 | ~ spl42_176 | spl42_223)), inference(subsumption_resolution, [], [f8720, f1260])).
fof(f8720, plain, (~ (e23 = op2(e22, e21)) | (~ spl42_11 | ~ spl42_176 | spl42_223)), inference(forward_demodulation, [], [f8719, f590])).
fof(f8719, plain, (~ (op2(e22, e21) = h12(e12)) | (~ spl42_11 | ~ spl42_176 | spl42_223)), inference(forward_demodulation, [], [f8718, f831])).
fof(f831, plain, ((e12 = op1(e13, e11)) | ~ spl42_11), inference(avatar_component_clause, [], [f829])).
fof(f829, plain, (spl42_11 <=> (e12 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_11])])).
fof(f8718, plain, (~ (op2(e22, e21) = h12(op1(e13, e11))) | (~ spl42_176 | spl42_223)), inference(forward_demodulation, [], [f2108, f1777])).
fof(f2108, plain, (~ (h12(op1(e13, e11)) = op2(e22, h12(e11))) | spl42_223), inference(avatar_component_clause, [], [f2106])).
fof(f2106, plain, (spl42_223 <=> (h12(op1(e13, e11)) = op2(e22, h12(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl42_223])])).
fof(f8662, plain, (~ spl42_41 | ~ spl42_125 | ~ spl42_176 | ~ spl42_180 | spl42_220), inference(avatar_contradiction_clause, [], [f8661])).
fof(f8661, plain, ($false | (~ spl42_41 | ~ spl42_125 | ~ spl42_176 | ~ spl42_180 | spl42_220)), inference(subsumption_resolution, [], [f8660, f1316])).
fof(f1316, plain, ((e20 = op2(e21, e21)) | ~ spl42_125), inference(avatar_component_clause, [], [f1314])).
fof(f1314, plain, (spl42_125 <=> (e20 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_125])])).
fof(f8660, plain, (~ (e20 = op2(e21, e21)) | (~ spl42_41 | ~ spl42_176 | ~ spl42_180 | spl42_220)), inference(forward_demodulation, [], [f8659, f6821])).
fof(f8659, plain, (~ (op2(e21, e21) = h12(e10)) | (~ spl42_41 | ~ spl42_176 | spl42_220)), inference(forward_demodulation, [], [f8658, f959])).
fof(f959, plain, ((e10 = op1(e11, e11)) | ~ spl42_41), inference(avatar_component_clause, [], [f957])).
fof(f957, plain, (spl42_41 <=> (e10 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_41])])).
fof(f8658, plain, (~ (op2(e21, e21) = h12(op1(e11, e11))) | (~ spl42_176 | spl42_220)), inference(forward_demodulation, [], [f2096, f1777])).
fof(f2096, plain, (~ (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) | spl42_220), inference(avatar_component_clause, [], [f2094])).
fof(f2094, plain, (spl42_220 <=> (h12(op1(e11, e11)) = op2(h12(e11), h12(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl42_220])])).
fof(f8606, plain, (~ spl42_28 | ~ spl42_95 | ~ spl42_176 | spl42_227), inference(avatar_contradiction_clause, [], [f8605])).
fof(f8605, plain, ($false | (~ spl42_28 | ~ spl42_95 | ~ spl42_176 | spl42_227)), inference(subsumption_resolution, [], [f8604, f1188])).
fof(f8604, plain, (~ (e22 = op2(e23, e21)) | (~ spl42_28 | ~ spl42_176 | spl42_227)), inference(forward_demodulation, [], [f8603, f591])).
fof(f8603, plain, (~ (op2(e23, e21) = h12(e13)) | (~ spl42_28 | ~ spl42_176 | spl42_227)), inference(forward_demodulation, [], [f8602, f903])).
fof(f903, plain, ((e13 = op1(e12, e11)) | ~ spl42_28), inference(avatar_component_clause, [], [f901])).
fof(f901, plain, (spl42_28 <=> (e13 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_28])])).
fof(f8602, plain, (~ (op2(e23, e21) = h12(op1(e12, e11))) | (~ spl42_176 | spl42_227)), inference(forward_demodulation, [], [f2124, f1777])).
fof(f2124, plain, (~ (h12(op1(e12, e11)) = op2(e23, h12(e11))) | spl42_227), inference(avatar_component_clause, [], [f2122])).
fof(f2122, plain, (spl42_227 <=> (h12(op1(e12, e11)) = op2(e23, h12(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl42_227])])).
fof(f8549, plain, (~ spl42_18 | spl42_225), inference(avatar_contradiction_clause, [], [f8548])).
fof(f8548, plain, ($false | (~ spl42_18 | spl42_225)), inference(trivial_inequality_removal, [], [f8547])).
fof(f8547, plain, (~ (h12(e11) = h12(e11)) | (~ spl42_18 | spl42_225)), inference(forward_demodulation, [], [f2116, f861])).
fof(f2116, plain, (~ (h12(e11) = h12(op1(e12, e13))) | spl42_225), inference(avatar_component_clause, [], [f2114])).
fof(f2114, plain, (spl42_225 <=> (h12(e11) = h12(op1(e12, e13)))), introduced(avatar_definition, [new_symbols(naming, [spl42_225])])).
fof(f8487, plain, (~ spl42_16 | ~ spl42_115 | ~ spl42_180 | spl42_224), inference(avatar_contradiction_clause, [], [f8486])).
fof(f8486, plain, ($false | (~ spl42_16 | ~ spl42_115 | ~ spl42_180 | spl42_224)), inference(subsumption_resolution, [], [f8485, f1273])).
fof(f1273, plain, ((e22 = op2(e22, e20)) | ~ spl42_115), inference(avatar_component_clause, [], [f1271])).
fof(f1271, plain, (spl42_115 <=> (e22 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_115])])).
fof(f8485, plain, (~ (e22 = op2(e22, e20)) | (~ spl42_16 | ~ spl42_180 | spl42_224)), inference(forward_demodulation, [], [f8484, f591])).
fof(f8484, plain, (~ (op2(e22, e20) = h12(e13)) | (~ spl42_16 | ~ spl42_180 | spl42_224)), inference(forward_demodulation, [], [f8483, f852])).
fof(f852, plain, ((e13 = op1(e13, e10)) | ~ spl42_16), inference(avatar_component_clause, [], [f850])).
fof(f850, plain, (spl42_16 <=> (e13 = op1(e13, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_16])])).
fof(f8483, plain, (~ (op2(e22, e20) = h12(op1(e13, e10))) | (~ spl42_180 | spl42_224)), inference(forward_demodulation, [], [f2112, f1800])).
fof(f2112, plain, (~ (h12(op1(e13, e10)) = op2(e22, h2(e10))) | spl42_224), inference(avatar_component_clause, [], [f2110])).
fof(f2110, plain, (spl42_224 <=> (h12(op1(e13, e10)) = op2(e22, h2(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl42_224])])).
fof(f8411, plain, (~ spl42_6 | ~ spl42_176 | ~ spl42_194 | spl42_222), inference(avatar_contradiction_clause, [], [f8410])).
fof(f8410, plain, ($false | (~ spl42_6 | ~ spl42_176 | ~ spl42_194 | spl42_222)), inference(subsumption_resolution, [], [f8398, f1777])).
fof(f8398, plain, (~ (e21 = h12(e11)) | (~ spl42_6 | ~ spl42_194 | spl42_222)), inference(backward_demodulation, [], [f8003, f810])).
fof(f810, plain, ((e11 = op1(e13, e12)) | ~ spl42_6), inference(avatar_component_clause, [], [f808])).
fof(f808, plain, (spl42_6 <=> (e11 = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_6])])).
fof(f8003, plain, (~ (e21 = h12(op1(e13, e12))) | (~ spl42_194 | spl42_222)), inference(forward_demodulation, [], [f2104, f1883])).
fof(f2104, plain, (~ (h9(e11) = h12(op1(e13, e12))) | spl42_222), inference(avatar_component_clause, [], [f2102])).
fof(f2102, plain, (spl42_222 <=> (h9(e11) = h12(op1(e13, e12)))), introduced(avatar_definition, [new_symbols(naming, [spl42_222])])).
fof(f8380, plain, (spl42_28 | ~ spl42_18 | ~ spl42_21 | ~ spl42_52), inference(avatar_split_clause, [], [f8365, f1003, f872, f859, f901])).
fof(f1003, plain, (spl42_52 <=> (e13 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_52])])).
fof(f8365, plain, ((e13 = op1(e12, e11)) | (~ spl42_18 | ~ spl42_21 | ~ spl42_52)), inference(backward_demodulation, [], [f5855, f1005])).
fof(f1005, plain, ((e13 = op1(e10, e13)) | ~ spl42_52), inference(avatar_component_clause, [], [f1003])).
fof(f5855, plain, ((op1(e10, e13) = op1(e12, e11)) | (~ spl42_18 | ~ spl42_21)), inference(forward_demodulation, [], [f5854, f874])).
fof(f5854, plain, ((op1(e12, e11) = op1(op1(e12, e12), e13)) | ~ spl42_18), inference(forward_demodulation, [], [f189, f861])).
fof(f189, plain, (op1(op1(e12, e12), e13) = op1(e12, op1(e12, e13))), inference(cnf_transformation, [], [f2])).
fof(f8303, plain, (spl42_16 | ~ spl42_65), inference(avatar_split_clause, [], [f8268, f1059, f850])).
fof(f1059, plain, (spl42_65 <=> (e10 = unit1)), introduced(avatar_definition, [new_symbols(naming, [spl42_65])])).
fof(f8268, plain, ((e13 = op1(e13, e10)) | ~ spl42_65), inference(backward_demodulation, [], [f217, f1061])).
fof(f1061, plain, ((e10 = unit1) | ~ spl42_65), inference(avatar_component_clause, [], [f1059])).
fof(f217, plain, (e13 = op1(e13, unit1)), inference(cnf_transformation, [], [f3])).
fof(f3, plain, (((e13 = unit1) | (e12 = unit1) | (e11 = unit1) | (e10 = unit1)) & (e13 = op1(e13, unit1)) & (e13 = op1(unit1, e13)) & (e12 = op1(e12, unit1)) & (e12 = op1(unit1, e12)) & (e11 = op1(e11, unit1)) & (e11 = op1(unit1, e11)) & (e10 = op1(e10, unit1)) & (e10 = op1(unit1, e10))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG015+1.p', ax3)).
fof(f8302, plain, (spl42_52 | ~ spl42_65), inference(avatar_split_clause, [], [f8267, f1059, f1003])).
fof(f8267, plain, ((e13 = op1(e10, e13)) | ~ spl42_65), inference(backward_demodulation, [], [f216, f1061])).
fof(f216, plain, (e13 = op1(unit1, e13)), inference(cnf_transformation, [], [f3])).
fof(f8301, plain, (spl42_31 | ~ spl42_65), inference(avatar_split_clause, [], [f8266, f1059, f914])).
fof(f8266, plain, ((e12 = op1(e12, e10)) | ~ spl42_65), inference(backward_demodulation, [], [f215, f1061])).
fof(f215, plain, (e12 = op1(e12, unit1)), inference(cnf_transformation, [], [f3])).
fof(f8299, plain, (spl42_58 | ~ spl42_65), inference(avatar_split_clause, [], [f8263, f1059, f1029])).
fof(f8263, plain, ((e11 = op1(e10, e11)) | ~ spl42_65), inference(backward_demodulation, [], [f212, f1061])).
fof(f212, plain, (e11 = op1(unit1, e11)), inference(cnf_transformation, [], [f3])).
fof(f8002, plain, (~ spl42_1 | ~ spl42_3), inference(avatar_contradiction_clause, [], [f8001])).
fof(f8001, plain, ($false | (~ spl42_1 | ~ spl42_3)), inference(subsumption_resolution, [], [f8000, f483])).
fof(f483, plain, ~ (e10 = e12), inference(cnf_transformation, [], [f19])).
fof(f8000, plain, ((e10 = e12) | (~ spl42_1 | ~ spl42_3)), inference(forward_demodulation, [], [f797, f789])).
fof(f797, plain, ((e12 = op1(e13, e13)) | ~ spl42_3), inference(avatar_component_clause, [], [f795])).
fof(f795, plain, (spl42_3 <=> (e12 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_3])])).
fof(f7395, plain, (~ spl42_1 | ~ spl42_180 | spl42_221), inference(avatar_contradiction_clause, [], [f7394])).
fof(f7394, plain, ($false | (~ spl42_1 | ~ spl42_180 | spl42_221)), inference(subsumption_resolution, [], [f7393, f1800])).
fof(f7393, plain, (~ (e20 = h2(e10)) | (~ spl42_1 | ~ spl42_180 | spl42_221)), inference(forward_demodulation, [], [f7392, f6821])).
fof(f7392, plain, (~ (h2(e10) = h12(e10)) | (~ spl42_1 | spl42_221)), inference(forward_demodulation, [], [f2100, f789])).
fof(f2100, plain, (~ (h2(e10) = h12(op1(e13, e13))) | spl42_221), inference(avatar_component_clause, [], [f2098])).
fof(f2098, plain, (spl42_221 <=> (h2(e10) = h12(op1(e13, e13)))), introduced(avatar_definition, [new_symbols(naming, [spl42_221])])).
fof(f7205, plain, (spl42_35 | ~ spl42_1 | ~ spl42_18 | ~ spl42_31), inference(avatar_split_clause, [], [f7204, f914, f859, f787, f931])).
fof(f7204, plain, ((e12 = op1(e11, e13)) | (~ spl42_1 | ~ spl42_18 | ~ spl42_31)), inference(forward_demodulation, [], [f5835, f916])).
fof(f7058, plain, (spl42_6 | ~ spl42_1 | ~ spl42_18 | ~ spl42_28 | ~ spl42_35 | ~ spl42_41 | ~ spl42_46), inference(avatar_split_clause, [], [f7049, f978, f957, f931, f901, f859, f787, f808])).
fof(f7049, plain, ((e11 = op1(e13, e12)) | (~ spl42_1 | ~ spl42_18 | ~ spl42_28 | ~ spl42_35 | ~ spl42_41 | ~ spl42_46)), inference(backward_demodulation, [], [f6567, f980])).
fof(f6567, plain, ((op1(e11, e10) = op1(e13, e12)) | (~ spl42_1 | ~ spl42_18 | ~ spl42_28 | ~ spl42_35 | ~ spl42_41)), inference(backward_demodulation, [], [f6507, f6566])).
fof(f6566, plain, ((op1(e11, e10) = op1(e12, op1(e10, e13))) | (~ spl42_1 | ~ spl42_18)), inference(forward_demodulation, [], [f4472, f5825])).
fof(f5825, plain, ((op1(e10, e13) = op1(e13, e10)) | ~ spl42_1), inference(forward_demodulation, [], [f209, f789])).
fof(f209, plain, (op1(op1(e13, e13), e13) = op1(e13, op1(e13, e13))), inference(cnf_transformation, [], [f2])).
fof(f4472, plain, ((op1(e11, e10) = op1(e12, op1(e13, e10))) | ~ spl42_18), inference(backward_demodulation, [], [f190, f861])).
fof(f190, plain, (op1(op1(e12, e13), e10) = op1(e12, op1(e13, e10))), inference(cnf_transformation, [], [f2])).
fof(f6507, plain, ((op1(e13, e12) = op1(e12, op1(e10, e13))) | (~ spl42_28 | ~ spl42_35 | ~ spl42_41)), inference(forward_demodulation, [], [f181, f5903])).
fof(f5903, plain, ((op1(e13, e12) = op1(op1(e12, e10), e13)) | (~ spl42_28 | ~ spl42_35 | ~ spl42_41)), inference(backward_demodulation, [], [f5795, f5894])).
fof(f5894, plain, ((op1(e12, e10) = op1(e13, e11)) | (~ spl42_28 | ~ spl42_41)), inference(forward_demodulation, [], [f5893, f903])).
fof(f5893, plain, ((op1(e12, e10) = op1(op1(e12, e11), e11)) | ~ spl42_41), inference(forward_demodulation, [], [f183, f959])).
fof(f183, plain, (op1(op1(e12, e11), e11) = op1(e12, op1(e11, e11))), inference(cnf_transformation, [], [f2])).
fof(f5795, plain, ((op1(e13, e12) = op1(op1(e13, e11), e13)) | ~ spl42_35), inference(forward_demodulation, [], [f201, f933])).
fof(f201, plain, (op1(op1(e13, e11), e13) = op1(e13, op1(e11, e13))), inference(cnf_transformation, [], [f2])).
fof(f181, plain, (op1(op1(e12, e10), e13) = op1(e12, op1(e10, e13))), inference(cnf_transformation, [], [f2])).
fof(f7057, plain, (~ spl42_1 | ~ spl42_28 | ~ spl42_35 | spl42_40 | ~ spl42_41 | ~ spl42_46), inference(avatar_contradiction_clause, [], [f7056])).
fof(f7056, plain, ($false | (~ spl42_1 | ~ spl42_28 | ~ spl42_35 | spl42_40 | ~ spl42_41 | ~ spl42_46)), inference(subsumption_resolution, [], [f7048, f953])).
fof(f953, plain, (~ (e13 = op1(e11, e12)) | spl42_40), inference(avatar_component_clause, [], [f952])).
fof(f7048, plain, ((e13 = op1(e11, e12)) | (~ spl42_1 | ~ spl42_28 | ~ spl42_35 | ~ spl42_41 | ~ spl42_46)), inference(backward_demodulation, [], [f6545, f980])).
fof(f6545, plain, ((e13 = op1(op1(e11, e10), e12)) | (~ spl42_1 | ~ spl42_28 | ~ spl42_35 | ~ spl42_41)), inference(forward_demodulation, [], [f164, f6524])).
fof(f6524, plain, ((e13 = op1(e11, op1(e10, e12))) | (~ spl42_1 | ~ spl42_28 | ~ spl42_35 | ~ spl42_41)), inference(backward_demodulation, [], [f5902, f6520])).
fof(f6520, plain, ((op1(e10, e12) = op1(e12, e10)) | (~ spl42_1 | ~ spl42_35 | ~ spl42_41)), inference(backward_demodulation, [], [f6497, f6519])).
fof(f6519, plain, ((op1(e10, e12) = op1(e11, op1(e10, e13))) | (~ spl42_35 | ~ spl42_41)), inference(forward_demodulation, [], [f5922, f5919])).
fof(f5919, plain, ((op1(e10, e13) = op1(e11, e12)) | (~ spl42_35 | ~ spl42_41)), inference(forward_demodulation, [], [f5918, f959])).
fof(f5918, plain, ((op1(e11, e12) = op1(op1(e11, e11), e13)) | ~ spl42_35), inference(forward_demodulation, [], [f169, f933])).
fof(f169, plain, (op1(op1(e11, e11), e13) = op1(e11, op1(e11, e13))), inference(cnf_transformation, [], [f2])).
fof(f5922, plain, ((op1(e10, e12) = op1(e11, op1(e11, e12))) | ~ spl42_41), inference(forward_demodulation, [], [f168, f959])).
fof(f168, plain, (op1(op1(e11, e11), e12) = op1(e11, op1(e11, e12))), inference(cnf_transformation, [], [f2])).
fof(f6497, plain, ((op1(e12, e10) = op1(e11, op1(e10, e13))) | (~ spl42_1 | ~ spl42_35)), inference(backward_demodulation, [], [f5805, f5825])).
fof(f5805, plain, ((op1(e12, e10) = op1(e11, op1(e13, e10))) | ~ spl42_35), inference(forward_demodulation, [], [f174, f933])).
fof(f174, plain, (op1(op1(e11, e13), e10) = op1(e11, op1(e13, e10))), inference(cnf_transformation, [], [f2])).
fof(f5902, plain, ((e13 = op1(e11, op1(e12, e10))) | (~ spl42_28 | ~ spl42_35 | ~ spl42_41)), inference(backward_demodulation, [], [f5792, f5894])).
fof(f5792, plain, ((e13 = op1(e11, op1(e13, e11))) | (~ spl42_28 | ~ spl42_35)), inference(forward_demodulation, [], [f5791, f903])).
fof(f5791, plain, ((op1(e12, e11) = op1(e11, op1(e13, e11))) | ~ spl42_35), inference(forward_demodulation, [], [f175, f933])).
fof(f175, plain, (op1(op1(e11, e13), e11) = op1(e11, op1(e13, e11))), inference(cnf_transformation, [], [f2])).
fof(f164, plain, (op1(op1(e11, e10), e12) = op1(e11, op1(e10, e12))), inference(cnf_transformation, [], [f2])).
fof(f6983, plain, (spl42_49 | ~ spl42_68), inference(avatar_split_clause, [], [f6964, f1071, f991])).
fof(f991, plain, (spl42_49 <=> (e10 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_49])])).
fof(f1071, plain, (spl42_68 <=> (e13 = unit1)), introduced(avatar_definition, [new_symbols(naming, [spl42_68])])).
fof(f6964, plain, ((e10 = op1(e10, e13)) | ~ spl42_68), inference(backward_demodulation, [], [f211, f1073])).
fof(f1073, plain, ((e13 = unit1) | ~ spl42_68), inference(avatar_component_clause, [], [f1071])).
fof(f211, plain, (e10 = op1(e10, unit1)), inference(cnf_transformation, [], [f3])).
fof(f6806, plain, (~ spl42_109 | ~ spl42_186), inference(avatar_split_clause, [], [f6791, f1834, f1246])).
fof(f1246, plain, (spl42_109 <=> (e20 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_109])])).
fof(f1834, plain, (spl42_186 <=> (e20 = h1(e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_186])])).
fof(f6791, plain, (~ (e20 = op2(e22, e21)) | ~ spl42_186), inference(backward_demodulation, [], [f1604, f1835])).
fof(f1835, plain, ((e20 = h1(e10)) | ~ spl42_186), inference(avatar_component_clause, [], [f1834])).
fof(f1604, plain, ~ (op2(e22, e21) = h1(e10)), inference(backward_demodulation, [], [f443, f548])).
fof(f548, plain, (op2(e21, e21) = h1(e10)), inference(cnf_transformation, [], [f26])).
fof(f26, plain, ((op2(e20, e21) = h1(e11)) & (op2(e21, e21) = h1(e10)) & (e21 = h1(e13)) & (e20 = h1(e12))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG015+1.p', ax26)).
fof(f443, plain, ~ (op2(e21, e21) = op2(e22, e21)), inference(cnf_transformation, [], [f18])).
fof(f18, plain, (~ (op2(e23, e22) = op2(e23, e23)) & ~ (op2(e23, e21) = op2(e23, e23)) & ~ (op2(e23, e21) = op2(e23, e22)) & ~ (op2(e23, e20) = op2(e23, e23)) & ~ (op2(e23, e20) = op2(e23, e22)) & ~ (op2(e23, e20) = op2(e23, e21)) & ~ (op2(e22, e22) = op2(e22, e23)) & ~ (op2(e22, e21) = op2(e22, e23)) & ~ (op2(e22, e21) = op2(e22, e22)) & ~ (op2(e22, e20) = op2(e22, e23)) & ~ (op2(e22, e20) = op2(e22, e22)) & ~ (op2(e22, e20) = op2(e22, e21)) & ~ (op2(e21, e22) = op2(e21, e23)) & ~ (op2(e21, e21) = op2(e21, e23)) & ~ (op2(e21, e21) = op2(e21, e22)) & ~ (op2(e21, e20) = op2(e21, e23)) & ~ (op2(e21, e20) = op2(e21, e22)) & ~ (op2(e21, e20) = op2(e21, e21)) & ~ (op2(e20, e22) = op2(e20, e23)) & ~ (op2(e20, e21) = op2(e20, e23)) & ~ (op2(e20, e21) = op2(e20, e22)) & ~ (op2(e20, e20) = op2(e20, e23)) & ~ (op2(e20, e20) = op2(e20, e22)) & ~ (op2(e20, e20) = op2(e20, e21)) & ~ (op2(e22, e23) = op2(e23, e23)) & ~ (op2(e21, e23) = op2(e23, e23)) & ~ (op2(e21, e23) = op2(e22, e23)) & ~ (op2(e20, e23) = op2(e23, e23)) & ~ (op2(e20, e23) = op2(e22, e23)) & ~ (op2(e20, e23) = op2(e21, e23)) & ~ (op2(e22, e22) = op2(e23, e22)) & ~ (op2(e21, e22) = op2(e23, e22)) & ~ (op2(e21, e22) = op2(e22, e22)) & ~ (op2(e20, e22) = op2(e23, e22)) & ~ (op2(e20, e22) = op2(e22, e22)) & ~ (op2(e20, e22) = op2(e21, e22)) & ~ (op2(e22, e21) = op2(e23, e21)) & ~ (op2(e21, e21) = op2(e23, e21)) & ~ (op2(e21, e21) = op2(e22, e21)) & ~ (op2(e20, e21) = op2(e23, e21)) & ~ (op2(e20, e21) = op2(e22, e21)) & ~ (op2(e20, e21) = op2(e21, e21)) & ~ (op2(e22, e20) = op2(e23, e20)) & ~ (op2(e21, e20) = op2(e23, e20)) & ~ (op2(e21, e20) = op2(e22, e20)) & ~ (op2(e20, e20) = op2(e23, e20)) & ~ (op2(e20, e20) = op2(e22, e20)) & ~ (op2(e20, e20) = op2(e21, e20))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG015+1.p', ax18)).
fof(f6721, plain, (~ spl42_96 | ~ spl42_198 | ~ spl42_393), inference(avatar_split_clause, [], [f6713, f3557, f1905, f1190])).
fof(f1190, plain, (spl42_96 <=> (e23 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_96])])).
fof(f3557, plain, (spl42_393 <=> (e23 = h3(e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_393])])).
fof(f6713, plain, (~ (e23 = op2(e23, e21)) | (~ spl42_198 | ~ spl42_393)), inference(backward_demodulation, [], [f6590, f3558])).
fof(f3558, plain, ((e23 = h3(e11)) | ~ spl42_393), inference(avatar_component_clause, [], [f3557])).
fof(f6590, plain, (~ (op2(e23, e21) = h3(e11)) | ~ spl42_198), inference(forward_demodulation, [], [f1738, f6000])).
fof(f6000, plain, ((h3(e11) = h10(e11)) | ~ spl42_198), inference(backward_demodulation, [], [f585, f5532])).
fof(f5532, plain, ((op2(e23, e20) = h3(e11)) | ~ spl42_198), inference(forward_demodulation, [], [f5518, f557])).
fof(f557, plain, (op2(e20, e23) = h3(e11)), inference(cnf_transformation, [], [f28])).
fof(f28, plain, ((op2(e20, e23) = h3(e11)) & (op2(e23, e23) = h3(e10)) & (e23 = h3(e13)) & (e20 = h3(e12))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG015+1.p', ax28)).
fof(f5518, plain, ((op2(e20, e23) = op2(e23, e20)) | ~ spl42_198), inference(backward_demodulation, [], [f1654, f1906])).
fof(f1654, plain, (op2(h3(e10), e23) = op2(e23, h3(e10))), inference(backward_demodulation, [], [f310, f556])).
fof(f556, plain, (op2(e23, e23) = h3(e10)), inference(cnf_transformation, [], [f28])).
fof(f310, plain, (op2(op2(e23, e23), e23) = op2(e23, op2(e23, e23))), inference(cnf_transformation, [], [f6])).
fof(f6, plain, ((op2(op2(e23, e23), e23) = op2(e23, op2(e23, e23))) & (op2(op2(e23, e23), e22) = op2(e23, op2(e23, e22))) & (op2(op2(e23, e23), e21) = op2(e23, op2(e23, e21))) & (op2(op2(e23, e23), e20) = op2(e23, op2(e23, e20))) & (op2(op2(e23, e22), e23) = op2(e23, op2(e22, e23))) & (op2(op2(e23, e22), e22) = op2(e23, op2(e22, e22))) & (op2(op2(e23, e22), e21) = op2(e23, op2(e22, e21))) & (op2(op2(e23, e22), e20) = op2(e23, op2(e22, e20))) & (op2(op2(e23, e21), e23) = op2(e23, op2(e21, e23))) & (op2(op2(e23, e21), e22) = op2(e23, op2(e21, e22))) & (op2(op2(e23, e21), e21) = op2(e23, op2(e21, e21))) & (op2(op2(e23, e21), e20) = op2(e23, op2(e21, e20))) & (op2(op2(e23, e20), e23) = op2(e23, op2(e20, e23))) & (op2(op2(e23, e20), e22) = op2(e23, op2(e20, e22))) & (op2(op2(e23, e20), e21) = op2(e23, op2(e20, e21))) & (op2(op2(e23, e20), e20) = op2(e23, op2(e20, e20))) & (op2(op2(e22, e23), e23) = op2(e22, op2(e23, e23))) & (op2(op2(e22, e23), e22) = op2(e22, op2(e23, e22))) & (op2(op2(e22, e23), e21) = op2(e22, op2(e23, e21))) & (op2(op2(e22, e23), e20) = op2(e22, op2(e23, e20))) & (op2(op2(e22, e22), e23) = op2(e22, op2(e22, e23))) & (op2(op2(e22, e22), e22) = op2(e22, op2(e22, e22))) & (op2(op2(e22, e22), e21) = op2(e22, op2(e22, e21))) & (op2(op2(e22, e22), e20) = op2(e22, op2(e22, e20))) & (op2(op2(e22, e21), e23) = op2(e22, op2(e21, e23))) & (op2(op2(e22, e21), e22) = op2(e22, op2(e21, e22))) & (op2(op2(e22, e21), e21) = op2(e22, op2(e21, e21))) & (op2(op2(e22, e21), e20) = op2(e22, op2(e21, e20))) & (op2(op2(e22, e20), e23) = op2(e22, op2(e20, e23))) & (op2(op2(e22, e20), e22) = op2(e22, op2(e20, e22))) & (op2(op2(e22, e20), e21) = op2(e22, op2(e20, e21))) & (op2(op2(e22, e20), e20) = op2(e22, op2(e20, e20))) & (op2(op2(e21, e23), e23) = op2(e21, op2(e23, e23))) & (op2(op2(e21, e23), e22) = op2(e21, op2(e23, e22))) & (op2(op2(e21, e23), e21) = op2(e21, op2(e23, e21))) & (op2(op2(e21, e23), e20) = op2(e21, op2(e23, e20))) & (op2(op2(e21, e22), e23) = op2(e21, op2(e22, e23))) & (op2(op2(e21, e22), e22) = op2(e21, op2(e22, e22))) & (op2(op2(e21, e22), e21) = op2(e21, op2(e22, e21))) & (op2(op2(e21, e22), e20) = op2(e21, op2(e22, e20))) & (op2(op2(e21, e21), e23) = op2(e21, op2(e21, e23))) & (op2(op2(e21, e21), e22) = op2(e21, op2(e21, e22))) & (op2(op2(e21, e21), e21) = op2(e21, op2(e21, e21))) & (op2(op2(e21, e21), e20) = op2(e21, op2(e21, e20))) & (op2(op2(e21, e20), e23) = op2(e21, op2(e20, e23))) & (op2(op2(e21, e20), e22) = op2(e21, op2(e20, e22))) & (op2(op2(e21, e20), e21) = op2(e21, op2(e20, e21))) & (op2(op2(e21, e20), e20) = op2(e21, op2(e20, e20))) & (op2(op2(e20, e23), e23) = op2(e20, op2(e23, e23))) & (op2(op2(e20, e23), e22) = op2(e20, op2(e23, e22))) & (op2(op2(e20, e23), e21) = op2(e20, op2(e23, e21))) & (op2(op2(e20, e23), e20) = op2(e20, op2(e23, e20))) & (op2(op2(e20, e22), e23) = op2(e20, op2(e22, e23))) & (op2(op2(e20, e22), e22) = op2(e20, op2(e22, e22))) & (op2(op2(e20, e22), e21) = op2(e20, op2(e22, e21))) & (op2(op2(e20, e22), e20) = op2(e20, op2(e22, e20))) & (op2(op2(e20, e21), e23) = op2(e20, op2(e21, e23))) & (op2(op2(e20, e21), e22) = op2(e20, op2(e21, e22))) & (op2(op2(e20, e21), e21) = op2(e20, op2(e21, e21))) & (op2(op2(e20, e21), e20) = op2(e20, op2(e21, e20))) & (op2(op2(e20, e20), e23) = op2(e20, op2(e20, e23))) & (op2(op2(e20, e20), e22) = op2(e20, op2(e20, e22))) & (op2(op2(e20, e20), e21) = op2(e20, op2(e20, e21))) & (op2(op2(e20, e20), e20) = op2(e20, op2(e20, e20)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG015+1.p', ax6)).
fof(f585, plain, (op2(e23, e20) = h10(e11)), inference(cnf_transformation, [], [f35])).
fof(f35, plain, ((op2(e23, e20) = h10(e11)) & (op2(e20, e20) = h10(e10)) & (e20 = h10(e13)) & (e23 = h10(e12))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG015+1.p', ax35)).
fof(f1738, plain, ~ (op2(e23, e21) = h10(e11)), inference(backward_demodulation, [], [f476, f585])).
fof(f476, plain, ~ (op2(e23, e20) = op2(e23, e21)), inference(cnf_transformation, [], [f18])).
fof(f6718, plain, (~ spl42_92 | ~ spl42_198 | ~ spl42_393), inference(avatar_split_clause, [], [f6706, f3557, f1905, f1173])).
fof(f1173, plain, (spl42_92 <=> (e23 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_92])])).
fof(f6706, plain, (~ (e23 = op2(e23, e22)) | (~ spl42_198 | ~ spl42_393)), inference(backward_demodulation, [], [f6001, f3558])).
fof(f6001, plain, (~ (op2(e23, e22) = h3(e11)) | ~ spl42_198), inference(backward_demodulation, [], [f1739, f6000])).
fof(f1739, plain, ~ (op2(e23, e22) = h10(e11)), inference(backward_demodulation, [], [f477, f585])).
fof(f477, plain, ~ (op2(e23, e20) = op2(e23, e22)), inference(cnf_transformation, [], [f18])).
fof(f6679, plain, (spl42_165 | ~ spl42_149), inference(avatar_split_clause, [], [f6221, f1416, f1484])).
fof(f1484, plain, (spl42_165 <=> (e20 = inv2(e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_165])])).
fof(f1416, plain, (spl42_149 <=> (e20 = unit2)), introduced(avatar_definition, [new_symbols(naming, [spl42_149])])).
fof(f6221, plain, ((e20 = inv2(e20)) | ~ spl42_149), inference(backward_demodulation, [], [f359, f1418])).
fof(f1418, plain, ((e20 = unit2) | ~ spl42_149), inference(avatar_component_clause, [], [f1416])).
fof(f359, plain, (unit2 = inv2(unit2)), inference(cnf_transformation, [], [f13])).
fof(f13, plain, (unit2 = inv2(unit2)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG015+1.p', ax13)).
fof(f6674, plain, (spl42_109 | ~ spl42_149 | ~ spl42_163), inference(avatar_split_clause, [], [f6673, f1475, f1416, f1246])).
fof(f1475, plain, (spl42_163 <=> (e22 = inv2(e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_163])])).
fof(f6673, plain, ((e20 = op2(e22, e21)) | (~ spl42_149 | ~ spl42_163)), inference(forward_demodulation, [], [f6672, f1418])).
fof(f6672, plain, ((op2(e22, e21) = unit2) | ~ spl42_163), inference(forward_demodulation, [], [f323, f1477])).
fof(f1477, plain, ((e22 = inv2(e21)) | ~ spl42_163), inference(avatar_component_clause, [], [f1475])).
fof(f323, plain, (unit2 = op2(inv2(e21), e21)), inference(cnf_transformation, [], [f8])).
fof(f8, plain, (((e23 = inv2(e23)) | (e22 = inv2(e23)) | (e21 = inv2(e23)) | (e20 = inv2(e23))) & ((e23 = inv2(e22)) | (e22 = inv2(e22)) | (e21 = inv2(e22)) | (e20 = inv2(e22))) & ((e23 = inv2(e21)) | (e22 = inv2(e21)) | (e21 = inv2(e21)) | (e20 = inv2(e21))) & ((e23 = inv2(e20)) | (e22 = inv2(e20)) | (e21 = inv2(e20)) | (e20 = inv2(e20))) & (unit2 = op2(inv2(e23), e23)) & (unit2 = op2(e23, inv2(e23))) & (unit2 = op2(inv2(e22), e22)) & (unit2 = op2(e22, inv2(e22))) & (unit2 = op2(inv2(e21), e21)) & (unit2 = op2(e21, inv2(e21))) & (unit2 = op2(inv2(e20), e20)) & (unit2 = op2(e20, inv2(e20)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG015+1.p', ax8)).
fof(f6644, plain, (~ spl42_94 | ~ spl42_142), inference(avatar_split_clause, [], [f6634, f1386, f1182])).
fof(f1182, plain, (spl42_94 <=> (e21 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_94])])).
fof(f6634, plain, (~ (e21 = op2(e23, e21)) | ~ spl42_142), inference(backward_demodulation, [], [f1594, f6632])).
fof(f6632, plain, ((e21 = h1(e11)) | ~ spl42_142), inference(forward_demodulation, [], [f549, f1388])).
fof(f549, plain, (op2(e20, e21) = h1(e11)), inference(cnf_transformation, [], [f26])).
fof(f1594, plain, ~ (op2(e23, e21) = h1(e11)), inference(backward_demodulation, [], [f442, f549])).
fof(f442, plain, ~ (op2(e20, e21) = op2(e23, e21)), inference(cnf_transformation, [], [f18])).
fof(f6592, plain, (spl42_109 | ~ spl42_149 | ~ spl42_158), inference(avatar_split_clause, [], [f6591, f1454, f1416, f1246])).
fof(f1454, plain, (spl42_158 <=> (e21 = inv2(e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_158])])).
fof(f6591, plain, ((e20 = op2(e22, e21)) | (~ spl42_149 | ~ spl42_158)), inference(forward_demodulation, [], [f6203, f1418])).
fof(f6203, plain, ((op2(e22, e21) = unit2) | ~ spl42_158), inference(backward_demodulation, [], [f324, f1456])).
fof(f1456, plain, ((e21 = inv2(e22)) | ~ spl42_158), inference(avatar_component_clause, [], [f1454])).
fof(f324, plain, (unit2 = op2(e22, inv2(e22))), inference(cnf_transformation, [], [f8])).
fof(f6587, plain, (spl42_142 | ~ spl42_149), inference(avatar_split_clause, [], [f6586, f1416, f1386])).
fof(f6586, plain, ((e21 = op2(e20, e21)) | ~ spl42_149), inference(forward_demodulation, [], [f313, f1418])).
fof(f313, plain, (e21 = op2(unit2, e21)), inference(cnf_transformation, [], [f7])).
fof(f7, plain, (((e23 = unit2) | (e22 = unit2) | (e21 = unit2) | (e20 = unit2)) & (e23 = op2(e23, unit2)) & (e23 = op2(unit2, e23)) & (e22 = op2(e22, unit2)) & (e22 = op2(unit2, e22)) & (e21 = op2(e21, unit2)) & (e21 = op2(unit2, e21)) & (e20 = op2(e20, unit2)) & (e20 = op2(unit2, e20))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG015+1.p', ax7)).
fof(f6505, plain, (~ spl42_49 | ~ spl42_1), inference(avatar_split_clause, [], [f5822, f787, f991])).
fof(f5822, plain, (~ (e10 = op1(e10, e13)) | ~ spl42_1), inference(forward_demodulation, [], [f406, f789])).
fof(f406, plain, ~ (op1(e10, e13) = op1(e13, e13)), inference(cnf_transformation, [], [f17])).
fof(f17, plain, (~ (op1(e13, e12) = op1(e13, e13)) & ~ (op1(e13, e11) = op1(e13, e13)) & ~ (op1(e13, e11) = op1(e13, e12)) & ~ (op1(e13, e10) = op1(e13, e13)) & ~ (op1(e13, e10) = op1(e13, e12)) & ~ (op1(e13, e10) = op1(e13, e11)) & ~ (op1(e12, e12) = op1(e12, e13)) & ~ (op1(e12, e11) = op1(e12, e13)) & ~ (op1(e12, e11) = op1(e12, e12)) & ~ (op1(e12, e10) = op1(e12, e13)) & ~ (op1(e12, e10) = op1(e12, e12)) & ~ (op1(e12, e10) = op1(e12, e11)) & ~ (op1(e11, e12) = op1(e11, e13)) & ~ (op1(e11, e11) = op1(e11, e13)) & ~ (op1(e11, e11) = op1(e11, e12)) & ~ (op1(e11, e10) = op1(e11, e13)) & ~ (op1(e11, e10) = op1(e11, e12)) & ~ (op1(e11, e10) = op1(e11, e11)) & ~ (op1(e10, e12) = op1(e10, e13)) & ~ (op1(e10, e11) = op1(e10, e13)) & ~ (op1(e10, e11) = op1(e10, e12)) & ~ (op1(e10, e10) = op1(e10, e13)) & ~ (op1(e10, e10) = op1(e10, e12)) & ~ (op1(e10, e10) = op1(e10, e11)) & ~ (op1(e12, e13) = op1(e13, e13)) & ~ (op1(e11, e13) = op1(e13, e13)) & ~ (op1(e11, e13) = op1(e12, e13)) & ~ (op1(e10, e13) = op1(e13, e13)) & ~ (op1(e10, e13) = op1(e12, e13)) & ~ (op1(e10, e13) = op1(e11, e13)) & ~ (op1(e12, e12) = op1(e13, e12)) & ~ (op1(e11, e12) = op1(e13, e12)) & ~ (op1(e11, e12) = op1(e12, e12)) & ~ (op1(e10, e12) = op1(e13, e12)) & ~ (op1(e10, e12) = op1(e12, e12)) & ~ (op1(e10, e12) = op1(e11, e12)) & ~ (op1(e12, e11) = op1(e13, e11)) & ~ (op1(e11, e11) = op1(e13, e11)) & ~ (op1(e11, e11) = op1(e12, e11)) & ~ (op1(e10, e11) = op1(e13, e11)) & ~ (op1(e10, e11) = op1(e12, e11)) & ~ (op1(e10, e11) = op1(e11, e11)) & ~ (op1(e12, e10) = op1(e13, e10)) & ~ (op1(e11, e10) = op1(e13, e10)) & ~ (op1(e11, e10) = op1(e12, e10)) & ~ (op1(e10, e10) = op1(e13, e10)) & ~ (op1(e10, e10) = op1(e12, e10)) & ~ (op1(e10, e10) = op1(e11, e10))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG015+1.p', ax17)).
fof(f6402, plain, (~ spl42_18 | ~ spl42_67), inference(avatar_contradiction_clause, [], [f6401])).
fof(f6401, plain, ($false | (~ spl42_18 | ~ spl42_67)), inference(subsumption_resolution, [], [f6400, f486])).
fof(f486, plain, ~ (e11 = e13), inference(cnf_transformation, [], [f19])).
fof(f6400, plain, ((e11 = e13) | (~ spl42_18 | ~ spl42_67)), inference(forward_demodulation, [], [f6376, f861])).
fof(f6376, plain, ((e13 = op1(e12, e13)) | ~ spl42_67), inference(backward_demodulation, [], [f216, f1069])).
fof(f1069, plain, ((e12 = unit1) | ~ spl42_67), inference(avatar_component_clause, [], [f1067])).
fof(f1067, plain, (spl42_67 <=> (e12 = unit1)), introduced(avatar_definition, [new_symbols(naming, [spl42_67])])).
fof(f6263, plain, (spl42_393 | ~ spl42_100 | ~ spl42_198), inference(avatar_split_clause, [], [f6262, f1905, f1207, f3557])).
fof(f6262, plain, ((e23 = h3(e11)) | (~ spl42_100 | ~ spl42_198)), inference(backward_demodulation, [], [f5532, f1209])).
fof(f6234, plain, (spl42_130 | ~ spl42_149), inference(avatar_split_clause, [], [f6218, f1416, f1335])).
fof(f6218, plain, ((e21 = op2(e21, e20)) | ~ spl42_149), inference(backward_demodulation, [], [f314, f1418])).
fof(f314, plain, (e21 = op2(e21, unit2)), inference(cnf_transformation, [], [f7])).
fof(f6051, plain, (~ spl42_102 | ~ spl42_123 | ~ spl42_125 | ~ spl42_194), inference(avatar_contradiction_clause, [], [f6050])).
fof(f6050, plain, ($false | (~ spl42_102 | ~ spl42_123 | ~ spl42_125 | ~ spl42_194)), inference(subsumption_resolution, [], [f6049, f488])).
fof(f488, plain, ~ (e20 = e21), inference(cnf_transformation, [], [f20])).
fof(f20, plain, (~ (e22 = e23) & ~ (e21 = e23) & ~ (e21 = e22) & ~ (e20 = e23) & ~ (e20 = e22) & ~ (e20 = e21)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG015+1.p', ax20)).
fof(f6049, plain, ((e20 = e21) | (~ spl42_102 | ~ spl42_123 | ~ spl42_125 | ~ spl42_194)), inference(forward_demodulation, [], [f6048, f1316])).
fof(f6048, plain, ((e21 = op2(e21, e21)) | (~ spl42_102 | ~ spl42_123 | ~ spl42_194)), inference(forward_demodulation, [], [f6047, f1218])).
fof(f1218, plain, ((e21 = op2(e22, e23)) | ~ spl42_102), inference(avatar_component_clause, [], [f1216])).
fof(f1216, plain, (spl42_102 <=> (e21 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_102])])).
fof(f6047, plain, ((op2(e21, e21) = op2(e22, e23)) | (~ spl42_123 | ~ spl42_194)), inference(forward_demodulation, [], [f4617, f5981])).
fof(f5981, plain, ((e22 = h5(e11)) | ~ spl42_123), inference(forward_demodulation, [], [f565, f1307])).
fof(f1307, plain, ((e22 = op2(e21, e22)) | ~ spl42_123), inference(avatar_component_clause, [], [f1305])).
fof(f1305, plain, (spl42_123 <=> (e22 = op2(e21, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_123])])).
fof(f565, plain, (op2(e21, e22) = h5(e11)), inference(cnf_transformation, [], [f30])).
fof(f30, plain, ((op2(e21, e22) = h5(e11)) & (op2(e22, e22) = h5(e10)) & (e22 = h5(e13)) & (e21 = h5(e12))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG015+1.p', ax30)).
fof(f4617, plain, ((op2(e21, e21) = op2(h5(e11), e23)) | ~ spl42_194), inference(backward_demodulation, [], [f1736, f1883])).
fof(f1736, plain, (op2(h5(e11), e23) = op2(e21, h9(e11))), inference(backward_demodulation, [], [f1687, f581])).
fof(f581, plain, (op2(e22, e23) = h9(e11)), inference(cnf_transformation, [], [f34])).
fof(f34, plain, ((op2(e22, e23) = h9(e11)) & (op2(e23, e23) = h9(e10)) & (e23 = h9(e13)) & (e22 = h9(e12))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG015+1.p', ax34)).
fof(f1687, plain, (op2(e21, op2(e22, e23)) = op2(h5(e11), e23)), inference(backward_demodulation, [], [f274, f565])).
fof(f274, plain, (op2(op2(e21, e22), e23) = op2(e21, op2(e22, e23))), inference(cnf_transformation, [], [f6])).
fof(f6046, plain, (spl42_186 | ~ spl42_125), inference(avatar_split_clause, [], [f6045, f1314, f1834])).
fof(f6045, plain, ((e20 = h1(e10)) | ~ spl42_125), inference(forward_demodulation, [], [f548, f1316])).
fof(f6015, plain, (spl42_180 | ~ spl42_105), inference(avatar_split_clause, [], [f6014, f1229, f1799])).
fof(f1229, plain, (spl42_105 <=> (e20 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_105])])).
fof(f6014, plain, ((e20 = h2(e10)) | ~ spl42_105), inference(forward_demodulation, [], [f552, f1231])).
fof(f1231, plain, ((e20 = op2(e22, e22)) | ~ spl42_105), inference(avatar_component_clause, [], [f1229])).
fof(f5735, plain, (~ spl42_41 | ~ spl42_66), inference(avatar_contradiction_clause, [], [f5734])).
fof(f5734, plain, ($false | (~ spl42_41 | ~ spl42_66)), inference(subsumption_resolution, [], [f5733, f482])).
fof(f5733, plain, ((e10 = e11) | (~ spl42_41 | ~ spl42_66)), inference(forward_demodulation, [], [f5715, f959])).
fof(f5715, plain, ((e11 = op1(e11, e11)) | ~ spl42_66), inference(backward_demodulation, [], [f213, f1065])).
fof(f1065, plain, ((e11 = unit1) | ~ spl42_66), inference(avatar_component_clause, [], [f1063])).
fof(f1063, plain, (spl42_66 <=> (e11 = unit1)), introduced(avatar_definition, [new_symbols(naming, [spl42_66])])).
fof(f213, plain, (e11 = op1(e11, unit1)), inference(cnf_transformation, [], [f3])).
fof(f5663, plain, (~ spl42_85 | ~ spl42_152), inference(avatar_contradiction_clause, [], [f5662])).
fof(f5662, plain, ($false | (~ spl42_85 | ~ spl42_152)), inference(subsumption_resolution, [], [f5661, f490])).
fof(f490, plain, ~ (e20 = e23), inference(cnf_transformation, [], [f20])).
fof(f5661, plain, ((e20 = e23) | (~ spl42_85 | ~ spl42_152)), inference(forward_demodulation, [], [f5638, f1146])).
fof(f1146, plain, ((e20 = op2(e23, e23)) | ~ spl42_85), inference(avatar_component_clause, [], [f1144])).
fof(f1144, plain, (spl42_85 <=> (e20 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_85])])).
fof(f5638, plain, ((e23 = op2(e23, e23)) | ~ spl42_152), inference(backward_demodulation, [], [f318, f1430])).
fof(f1430, plain, ((e23 = unit2) | ~ spl42_152), inference(avatar_component_clause, [], [f1428])).
fof(f1428, plain, (spl42_152 <=> (e23 = unit2)), introduced(avatar_definition, [new_symbols(naming, [spl42_152])])).
fof(f318, plain, (e23 = op2(e23, unit2)), inference(cnf_transformation, [], [f7])).
fof(f5604, plain, (~ spl42_164 | ~ spl42_165 | ~ spl42_194 | ~ spl42_198), inference(avatar_contradiction_clause, [], [f5603])).
fof(f5603, plain, ($false | (~ spl42_164 | ~ spl42_165 | ~ spl42_194 | ~ spl42_198)), inference(subsumption_resolution, [], [f5586, f3847])).
fof(f3847, plain, (~ (op2(e22, e20) = unit2) | ~ spl42_165), inference(backward_demodulation, [], [f1678, f3842])).
fof(f3842, plain, ((unit2 = h4(e10)) | ~ spl42_165), inference(backward_demodulation, [], [f560, f3833])).
fof(f3833, plain, ((op2(e20, e20) = unit2) | ~ spl42_165), inference(backward_demodulation, [], [f320, f1486])).
fof(f1486, plain, ((e20 = inv2(e20)) | ~ spl42_165), inference(avatar_component_clause, [], [f1484])).
fof(f320, plain, (unit2 = op2(e20, inv2(e20))), inference(cnf_transformation, [], [f8])).
fof(f560, plain, (op2(e20, e20) = h4(e10)), inference(cnf_transformation, [], [f29])).
fof(f29, plain, ((op2(e21, e20) = h4(e11)) & (op2(e20, e20) = h4(e10)) & (e20 = h4(e13)) & (e21 = h4(e12))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG015+1.p', ax29)).
fof(f1678, plain, ~ (op2(e22, e20) = h4(e10)), inference(backward_demodulation, [], [f435, f560])).
fof(f435, plain, ~ (op2(e20, e20) = op2(e22, e20)), inference(cnf_transformation, [], [f18])).
fof(f5586, plain, ((op2(e22, e20) = unit2) | (~ spl42_164 | ~ spl42_194 | ~ spl42_198)), inference(backward_demodulation, [], [f5528, f5579])).
fof(f5579, plain, ((unit2 = h6(e11)) | ~ spl42_164), inference(backward_demodulation, [], [f569, f5572])).
fof(f5572, plain, ((op2(e21, e23) = unit2) | ~ spl42_164), inference(backward_demodulation, [], [f322, f1481])).
fof(f1481, plain, ((e23 = inv2(e21)) | ~ spl42_164), inference(avatar_component_clause, [], [f1479])).
fof(f1479, plain, (spl42_164 <=> (e23 = inv2(e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_164])])).
fof(f322, plain, (unit2 = op2(e21, inv2(e21))), inference(cnf_transformation, [], [f8])).
fof(f569, plain, (op2(e21, e23) = h6(e11)), inference(cnf_transformation, [], [f31])).
fof(f31, plain, ((op2(e21, e23) = h6(e11)) & (op2(e23, e23) = h6(e10)) & (e23 = h6(e13)) & (e21 = h6(e12))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG015+1.p', ax31)).
fof(f5528, plain, ((op2(e22, e20) = h6(e11)) | (~ spl42_194 | ~ spl42_198)), inference(backward_demodulation, [], [f5077, f1906])).
fof(f5077, plain, ((h6(e11) = op2(e22, h3(e10))) | ~ spl42_194), inference(forward_demodulation, [], [f5076, f569])).
fof(f5076, plain, ((op2(e21, e23) = op2(e22, h3(e10))) | ~ spl42_194), inference(forward_demodulation, [], [f1735, f1883])).
fof(f1735, plain, (op2(e22, h3(e10)) = op2(h9(e11), e23)), inference(backward_demodulation, [], [f1650, f581])).
fof(f1650, plain, (op2(op2(e22, e23), e23) = op2(e22, h3(e10))), inference(backward_demodulation, [], [f294, f556])).
fof(f294, plain, (op2(op2(e22, e23), e23) = op2(e22, op2(e23, e23))), inference(cnf_transformation, [], [f6])).
fof(f5467, plain, (spl42_94 | ~ spl42_91 | ~ spl42_102 | ~ spl42_194), inference(avatar_split_clause, [], [f5458, f1882, f1216, f1169, f1182])).
fof(f1169, plain, (spl42_91 <=> (e22 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_91])])).
fof(f5458, plain, ((e21 = op2(e23, e21)) | (~ spl42_91 | ~ spl42_102 | ~ spl42_194)), inference(backward_demodulation, [], [f589, f5457])).
fof(f5457, plain, ((e21 = h11(e11)) | (~ spl42_91 | ~ spl42_102 | ~ spl42_194)), inference(forward_demodulation, [], [f5404, f1218])).
fof(f5404, plain, ((op2(e22, e23) = h11(e11)) | (~ spl42_91 | ~ spl42_194)), inference(backward_demodulation, [], [f4630, f5397])).
fof(f5397, plain, ((e22 = h12(e11)) | ~ spl42_91), inference(forward_demodulation, [], [f593, f1171])).
fof(f1171, plain, ((e22 = op2(e23, e22)) | ~ spl42_91), inference(avatar_component_clause, [], [f1169])).
fof(f593, plain, (op2(e23, e22) = h12(e11)), inference(cnf_transformation, [], [f37])).
fof(f4630, plain, ((h11(e11) = op2(h12(e11), e23)) | ~ spl42_194), inference(forward_demodulation, [], [f4621, f589])).
fof(f4621, plain, ((op2(e23, e21) = op2(h12(e11), e23)) | ~ spl42_194), inference(backward_demodulation, [], [f1766, f1883])).
fof(f1766, plain, (op2(e23, h9(e11)) = op2(h12(e11), e23)), inference(backward_demodulation, [], [f1732, f593])).
fof(f1732, plain, (op2(op2(e23, e22), e23) = op2(e23, h9(e11))), inference(backward_demodulation, [], [f306, f581])).
fof(f306, plain, (op2(op2(e23, e22), e23) = op2(e23, op2(e22, e23))), inference(cnf_transformation, [], [f6])).
fof(f589, plain, (op2(e23, e21) = h11(e11)), inference(cnf_transformation, [], [f36])).
fof(f36, plain, ((op2(e23, e21) = h11(e11)) & (op2(e21, e21) = h11(e10)) & (e21 = h11(e13)) & (e23 = h11(e12))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG015+1.p', ax36)).
fof(f5370, plain, (spl42_177 | ~ spl42_106), inference(avatar_split_clause, [], [f5369, f1233, f1782])).
fof(f1782, plain, (spl42_177 <=> (e21 = h2(e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_177])])).
fof(f1233, plain, (spl42_106 <=> (e21 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_106])])).
fof(f5369, plain, ((e21 = h2(e10)) | ~ spl42_106), inference(forward_demodulation, [], [f552, f1235])).
fof(f1235, plain, ((e21 = op2(e22, e22)) | ~ spl42_106), inference(avatar_component_clause, [], [f1233])).
fof(f5309, plain, (spl42_11 | ~ spl42_6 | ~ spl42_18 | ~ spl42_35), inference(avatar_split_clause, [], [f5302, f931, f859, f808, f829])).
fof(f5302, plain, ((e12 = op1(e13, e11)) | (~ spl42_6 | ~ spl42_18 | ~ spl42_35)), inference(backward_demodulation, [], [f4955, f933])).
fof(f4955, plain, ((op1(e11, e13) = op1(e13, e11)) | (~ spl42_6 | ~ spl42_18)), inference(forward_demodulation, [], [f4475, f810])).
fof(f4475, plain, ((op1(e13, e11) = op1(op1(e13, e12), e13)) | ~ spl42_18), inference(backward_demodulation, [], [f205, f861])).
fof(f205, plain, (op1(op1(e13, e12), e13) = op1(e13, op1(e12, e13))), inference(cnf_transformation, [], [f2])).
fof(f5202, plain, (~ spl42_102 | ~ spl42_151), inference(avatar_contradiction_clause, [], [f5201])).
fof(f5201, plain, ($false | (~ spl42_102 | ~ spl42_151)), inference(subsumption_resolution, [], [f5200, f492])).
fof(f492, plain, ~ (e21 = e23), inference(cnf_transformation, [], [f20])).
fof(f5200, plain, ((e21 = e23) | (~ spl42_102 | ~ spl42_151)), inference(forward_demodulation, [], [f5172, f1218])).
fof(f5172, plain, ((e23 = op2(e22, e23)) | ~ spl42_151), inference(backward_demodulation, [], [f317, f1426])).
fof(f1426, plain, ((e22 = unit2) | ~ spl42_151), inference(avatar_component_clause, [], [f1424])).
fof(f1424, plain, (spl42_151 <=> (e22 = unit2)), introduced(avatar_definition, [new_symbols(naming, [spl42_151])])).
fof(f317, plain, (e23 = op2(unit2, e23)), inference(cnf_transformation, [], [f7])).
fof(f5136, plain, (spl42_150 | ~ spl42_102 | ~ spl42_160), inference(avatar_split_clause, [], [f5135, f1462, f1216, f1420])).
fof(f1420, plain, (spl42_150 <=> (e21 = unit2)), introduced(avatar_definition, [new_symbols(naming, [spl42_150])])).
fof(f1462, plain, (spl42_160 <=> (e23 = inv2(e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_160])])).
fof(f5135, plain, ((e21 = unit2) | (~ spl42_102 | ~ spl42_160)), inference(forward_demodulation, [], [f5131, f1218])).
fof(f5131, plain, ((op2(e22, e23) = unit2) | ~ spl42_160), inference(backward_demodulation, [], [f324, f1464])).
fof(f1464, plain, ((e23 = inv2(e22)) | ~ spl42_160), inference(avatar_component_clause, [], [f1462])).
fof(f5127, plain, (~ spl42_102 | ~ spl42_107 | ~ spl42_194), inference(avatar_contradiction_clause, [], [f5126])).
fof(f5126, plain, ($false | (~ spl42_102 | ~ spl42_107 | ~ spl42_194)), inference(subsumption_resolution, [], [f5125, f4087])).
fof(f4087, plain, (~ (e21 = h8(e11)) | ~ spl42_102), inference(backward_demodulation, [], [f1720, f1218])).
fof(f1720, plain, ~ (op2(e22, e23) = h8(e11)), inference(backward_demodulation, [], [f474, f577])).
fof(f577, plain, (op2(e22, e21) = h8(e11)), inference(cnf_transformation, [], [f33])).
fof(f33, plain, ((op2(e22, e21) = h8(e11)) & (op2(e21, e21) = h8(e10)) & (e21 = h8(e13)) & (e22 = h8(e12))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG015+1.p', ax33)).
fof(f474, plain, ~ (op2(e22, e21) = op2(e22, e23)), inference(cnf_transformation, [], [f18])).
fof(f5125, plain, ((e21 = h8(e11)) | (~ spl42_102 | ~ spl42_107 | ~ spl42_194)), inference(forward_demodulation, [], [f5112, f1218])).
fof(f5112, plain, ((op2(e22, e23) = h8(e11)) | (~ spl42_107 | ~ spl42_194)), inference(backward_demodulation, [], [f5088, f5100])).
fof(f5100, plain, ((e22 = h2(e10)) | ~ spl42_107), inference(forward_demodulation, [], [f552, f1239])).
fof(f1239, plain, ((e22 = op2(e22, e22)) | ~ spl42_107), inference(avatar_component_clause, [], [f1237])).
fof(f1237, plain, (spl42_107 <=> (e22 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_107])])).
fof(f5088, plain, ((h8(e11) = op2(h2(e10), e23)) | ~ spl42_194), inference(forward_demodulation, [], [f5087, f577])).
fof(f5087, plain, ((op2(e22, e21) = op2(h2(e10), e23)) | ~ spl42_194), inference(forward_demodulation, [], [f1734, f1883])).
fof(f1734, plain, (op2(h2(e10), e23) = op2(e22, h9(e11))), inference(backward_demodulation, [], [f1628, f581])).
fof(f1628, plain, (op2(e22, op2(e22, e23)) = op2(h2(e10), e23)), inference(backward_demodulation, [], [f290, f552])).
fof(f290, plain, (op2(op2(e22, e22), e23) = op2(e22, op2(e22, e23))), inference(cnf_transformation, [], [f6])).
fof(f5097, plain, (~ spl42_177 | ~ spl42_102), inference(avatar_split_clause, [], [f5096, f1216, f1782])).
fof(f5096, plain, (~ (e21 = h2(e10)) | ~ spl42_102), inference(forward_demodulation, [], [f1634, f1218])).
fof(f1634, plain, ~ (op2(e22, e23) = h2(e10)), inference(backward_demodulation, [], [f475, f552])).
fof(f475, plain, ~ (op2(e22, e22) = op2(e22, e23)), inference(cnf_transformation, [], [f18])).
fof(f5082, plain, (spl42_198 | ~ spl42_85), inference(avatar_split_clause, [], [f5081, f1144, f1905])).
fof(f5081, plain, ((e20 = h3(e10)) | ~ spl42_85), inference(forward_demodulation, [], [f556, f1146])).
fof(f4823, plain, (~ spl42_52 | ~ spl42_159 | spl42_232), inference(avatar_contradiction_clause, [], [f4822])).
fof(f4822, plain, ($false | (~ spl42_52 | ~ spl42_159 | spl42_232)), inference(subsumption_resolution, [], [f4816, f591])).
fof(f4816, plain, (~ (e22 = h12(e13)) | (~ spl42_52 | ~ spl42_159 | spl42_232)), inference(backward_demodulation, [], [f3933, f1005])).
fof(f3933, plain, (~ (e22 = h12(op1(e10, e13))) | (~ spl42_159 | spl42_232)), inference(forward_demodulation, [], [f3923, f316])).
fof(f316, plain, (e22 = op2(e22, unit2)), inference(cnf_transformation, [], [f7])).
fof(f3923, plain, (~ (op2(e22, unit2) = h12(op1(e10, e13))) | (~ spl42_159 | spl42_232)), inference(backward_demodulation, [], [f2144, f3907])).
fof(f3907, plain, ((unit2 = h2(e10)) | ~ spl42_159), inference(backward_demodulation, [], [f552, f3904])).
fof(f3904, plain, ((op2(e22, e22) = unit2) | ~ spl42_159), inference(backward_demodulation, [], [f324, f1460])).
fof(f1460, plain, ((e22 = inv2(e22)) | ~ spl42_159), inference(avatar_component_clause, [], [f1458])).
fof(f1458, plain, (spl42_159 <=> (e22 = inv2(e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_159])])).
fof(f2144, plain, (~ (h12(op1(e10, e13)) = op2(e22, h2(e10))) | spl42_232), inference(avatar_component_clause, [], [f2142])).
fof(f2142, plain, (spl42_232 <=> (h12(op1(e10, e13)) = op2(e22, h2(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl42_232])])).
fof(f4765, plain, (spl42_176 | ~ spl42_90), inference(avatar_split_clause, [], [f4762, f1165, f1776])).
fof(f1165, plain, (spl42_90 <=> (e21 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_90])])).
fof(f4762, plain, ((e21 = h12(e11)) | ~ spl42_90), inference(backward_demodulation, [], [f593, f1167])).
fof(f1167, plain, ((e21 = op2(e23, e22)) | ~ spl42_90), inference(avatar_component_clause, [], [f1165])).
fof(f4673, plain, (spl42_123 | ~ spl42_150), inference(avatar_split_clause, [], [f4639, f1420, f1305])).
fof(f4639, plain, ((e22 = op2(e21, e22)) | ~ spl42_150), inference(backward_demodulation, [], [f315, f1422])).
fof(f1422, plain, ((e21 = unit2) | ~ spl42_150), inference(avatar_component_clause, [], [f1420])).
fof(f315, plain, (e22 = op2(unit2, e22)), inference(cnf_transformation, [], [f7])).
fof(f4631, plain, (spl42_112 | ~ spl42_159 | ~ spl42_194), inference(avatar_split_clause, [], [f4622, f1882, f1458, f1258])).
fof(f4622, plain, ((e23 = op2(e22, e21)) | (~ spl42_159 | ~ spl42_194)), inference(backward_demodulation, [], [f3931, f1883])).
fof(f3931, plain, ((e23 = op2(e22, h9(e11))) | ~ spl42_159), inference(forward_demodulation, [], [f3920, f317])).
fof(f3920, plain, ((op2(unit2, e23) = op2(e22, h9(e11))) | ~ spl42_159), inference(backward_demodulation, [], [f1734, f3907])).
fof(f4550, plain, (~ spl42_1 | ~ spl42_4), inference(avatar_contradiction_clause, [], [f4549])).
fof(f4549, plain, ($false | (~ spl42_1 | ~ spl42_4)), inference(subsumption_resolution, [], [f4548, f484])).
fof(f484, plain, ~ (e10 = e13), inference(cnf_transformation, [], [f19])).
fof(f4548, plain, ((e10 = e13) | (~ spl42_1 | ~ spl42_4)), inference(backward_demodulation, [], [f801, f789])).
fof(f801, plain, ((e13 = op1(e13, e13)) | ~ spl42_4), inference(avatar_component_clause, [], [f799])).
fof(f799, plain, (spl42_4 <=> (e13 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_4])])).
fof(f4286, plain, (spl42_46 | ~ spl42_65), inference(avatar_split_clause, [], [f4251, f1059, f978])).
fof(f4251, plain, ((e11 = op1(e11, e10)) | ~ spl42_65), inference(backward_demodulation, [], [f213, f1061])).
fof(f4127, plain, (~ spl42_85 | ~ spl42_88), inference(avatar_contradiction_clause, [], [f4126])).
fof(f4126, plain, ($false | (~ spl42_85 | ~ spl42_88)), inference(subsumption_resolution, [], [f4125, f490])).
fof(f4125, plain, ((e20 = e23) | (~ spl42_85 | ~ spl42_88)), inference(backward_demodulation, [], [f1158, f1146])).
fof(f1158, plain, ((e23 = op2(e23, e23)) | ~ spl42_88), inference(avatar_component_clause, [], [f1156])).
fof(f1156, plain, (spl42_88 <=> (e23 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_88])])).
fof(f4088, plain, (spl42_194 | ~ spl42_102), inference(avatar_split_clause, [], [f4084, f1216, f1882])).
fof(f4084, plain, ((e21 = h9(e11)) | ~ spl42_102), inference(backward_demodulation, [], [f581, f1218])).
fof(f4023, plain, (~ spl42_89 | ~ spl42_149 | ~ spl42_159), inference(avatar_split_clause, [], [f3992, f1458, f1416, f1161])).
fof(f1161, plain, (spl42_89 <=> (e20 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_89])])).
fof(f3992, plain, (~ (e20 = op2(e23, e22)) | (~ spl42_149 | ~ spl42_159)), inference(backward_demodulation, [], [f3910, f1418])).
fof(f3910, plain, (~ (op2(e23, e22) = unit2) | ~ spl42_159), inference(backward_demodulation, [], [f1631, f3907])).
fof(f1631, plain, ~ (op2(e23, e22) = h2(e10)), inference(backward_demodulation, [], [f451, f552])).
fof(f451, plain, ~ (op2(e22, e22) = op2(e23, e22)), inference(cnf_transformation, [], [f18])).
fof(f4017, plain, (~ spl42_93 | ~ spl42_149 | ~ spl42_162), inference(avatar_split_clause, [], [f3984, f1471, f1416, f1178])).
fof(f1178, plain, (spl42_93 <=> (e20 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_93])])).
fof(f1471, plain, (spl42_162 <=> (e21 = inv2(e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_162])])).
fof(f3984, plain, (~ (e20 = op2(e23, e21)) | (~ spl42_149 | ~ spl42_162)), inference(backward_demodulation, [], [f3879, f1418])).
fof(f3879, plain, (~ (op2(e23, e21) = unit2) | ~ spl42_162), inference(backward_demodulation, [], [f1605, f3876])).
fof(f3876, plain, ((unit2 = h1(e10)) | ~ spl42_162), inference(backward_demodulation, [], [f548, f3869])).
fof(f3869, plain, ((op2(e21, e21) = unit2) | ~ spl42_162), inference(backward_demodulation, [], [f322, f1473])).
fof(f1473, plain, ((e21 = inv2(e21)) | ~ spl42_162), inference(avatar_component_clause, [], [f1471])).
fof(f1605, plain, ~ (op2(e23, e21) = h1(e10)), inference(backward_demodulation, [], [f444, f548])).
fof(f444, plain, ~ (op2(e21, e21) = op2(e23, e21)), inference(cnf_transformation, [], [f18])).
fof(f4010, plain, (spl42_100 | ~ spl42_149), inference(avatar_split_clause, [], [f3968, f1416, f1207])).
fof(f3968, plain, ((e23 = op2(e23, e20)) | ~ spl42_149), inference(backward_demodulation, [], [f318, f1418])).
fof(f4008, plain, (spl42_115 | ~ spl42_149), inference(avatar_split_clause, [], [f3966, f1416, f1271])).
fof(f3966, plain, ((e22 = op2(e22, e20)) | ~ spl42_149), inference(backward_demodulation, [], [f316, f1418])).
fof(f3867, plain, (~ spl42_157 | ~ spl42_165), inference(avatar_split_clause, [], [f3837, f1484, f1450])).
fof(f1450, plain, (spl42_157 <=> (e20 = inv2(e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_157])])).
fof(f3837, plain, (~ (e20 = inv2(e22)) | ~ spl42_165), inference(backward_demodulation, [], [f381, f1486])).
fof(f381, plain, ~ (inv2(e20) = inv2(e22)), inference(cnf_transformation, [], [f16])).
fof(f16, plain, (~ (inv2(e22) = inv2(e23)) & ~ (inv2(e21) = inv2(e23)) & ~ (inv2(e21) = inv2(e22)) & ~ (inv2(e20) = inv2(e23)) & ~ (inv2(e20) = inv2(e22)) & ~ (inv2(e20) = inv2(e21))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG015+1.p', ax16)).
fof(f3866, plain, (~ spl42_161 | ~ spl42_165), inference(avatar_split_clause, [], [f3836, f1484, f1467])).
fof(f1467, plain, (spl42_161 <=> (e20 = inv2(e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_161])])).
fof(f3836, plain, (~ (e20 = inv2(e21)) | ~ spl42_165), inference(backward_demodulation, [], [f380, f1486])).
fof(f380, plain, ~ (inv2(e20) = inv2(e21)), inference(cnf_transformation, [], [f16])).
fof(f2157, plain, (~ spl42_220 | spl42_178 | spl42_175 | ~ spl42_221 | ~ spl42_222 | ~ spl42_223 | ~ spl42_224 | ~ spl42_225 | ~ spl42_226 | ~ spl42_227 | ~ spl42_228 | ~ spl42_229 | ~ spl42_230 | ~ spl42_231 | ~ spl42_232 | ~ spl42_233 | ~ spl42_234 | ~ spl42_235), inference(avatar_split_clause, [], [f2092, f2154, f2150, f2146, f2142, f2138, f2134, f2130, f2126, f2122, f2118, f2114, f2110, f2106, f2102, f2098, f1772, f1789, f2094])).
fof(f1789, plain, (spl42_178 <=> sP39), introduced(avatar_definition, [new_symbols(naming, [spl42_178])])).
fof(f1772, plain, (spl42_175 <=> sP40), introduced(avatar_definition, [new_symbols(naming, [spl42_175])])).
fof(f2092, plain, (~ (op2(h2(e10), h2(e10)) = h12(op1(e10, e10))) | ~ (h12(op1(e10, e11)) = op2(h2(e10), h12(e11))) | ~ (h12(op1(e10, e12)) = op2(e22, h9(e11))) | ~ (h12(op1(e10, e13)) = op2(e22, h2(e10))) | ~ (h12(op1(e11, e10)) = op2(h12(e11), h2(e10))) | ~ (h12(op1(e11, e12)) = op2(e23, h9(e11))) | ~ (h12(op1(e11, e13)) = op2(e23, h2(e10))) | ~ (h12(op1(e12, e10)) = op2(e23, h2(e10))) | ~ (h12(op1(e12, e11)) = op2(e23, h12(e11))) | ~ (h3(e10) = h12(op1(e12, e12))) | ~ (h12(e11) = h12(op1(e12, e13))) | ~ (h12(op1(e13, e10)) = op2(e22, h2(e10))) | ~ (h12(op1(e13, e11)) = op2(e22, h12(e11))) | ~ (h9(e11) = h12(op1(e13, e12))) | ~ (h2(e10) = h12(op1(e13, e13))) | sP40 | sP39 | ~ (h12(op1(e11, e11)) = op2(h12(e11), h12(e11)))), inference(forward_demodulation, [], [f2091, f1767])).
fof(f2091, plain, (~ (h12(op1(e10, e11)) = op2(h2(e10), h12(e11))) | ~ (h12(op1(e10, e12)) = op2(e22, h9(e11))) | ~ (h12(op1(e10, e13)) = op2(e22, h2(e10))) | ~ (h12(op1(e11, e10)) = op2(h12(e11), h2(e10))) | ~ (h12(op1(e11, e12)) = op2(e23, h9(e11))) | ~ (h12(op1(e11, e13)) = op2(e23, h2(e10))) | ~ (h12(op1(e12, e10)) = op2(e23, h2(e10))) | ~ (h12(op1(e12, e11)) = op2(e23, h12(e11))) | ~ (h3(e10) = h12(op1(e12, e12))) | ~ (h12(e11) = h12(op1(e12, e13))) | ~ (h12(op1(e13, e10)) = op2(e22, h2(e10))) | ~ (h12(op1(e13, e11)) = op2(e22, h12(e11))) | ~ (h9(e11) = h12(op1(e13, e12))) | ~ (h2(e10) = h12(op1(e13, e13))) | sP40 | sP39 | ~ (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) | ~ (h12(op1(e10, e10)) = op2(h12(e10), h12(e10)))), inference(forward_demodulation, [], [f2090, f1767])).
fof(f2090, plain, (~ (h12(op1(e10, e12)) = op2(e22, h9(e11))) | ~ (h12(op1(e10, e13)) = op2(e22, h2(e10))) | ~ (h12(op1(e11, e10)) = op2(h12(e11), h2(e10))) | ~ (h12(op1(e11, e12)) = op2(e23, h9(e11))) | ~ (h12(op1(e11, e13)) = op2(e23, h2(e10))) | ~ (h12(op1(e12, e10)) = op2(e23, h2(e10))) | ~ (h12(op1(e12, e11)) = op2(e23, h12(e11))) | ~ (h3(e10) = h12(op1(e12, e12))) | ~ (h12(e11) = h12(op1(e12, e13))) | ~ (h12(op1(e13, e10)) = op2(e22, h2(e10))) | ~ (h12(op1(e13, e11)) = op2(e22, h12(e11))) | ~ (h9(e11) = h12(op1(e13, e12))) | ~ (h2(e10) = h12(op1(e13, e13))) | sP40 | sP39 | ~ (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) | ~ (h12(op1(e10, e11)) = op2(h12(e10), h12(e11))) | ~ (h12(op1(e10, e10)) = op2(h12(e10), h12(e10)))), inference(forward_demodulation, [], [f2089, f1734])).
fof(f2089, plain, (~ (h12(op1(e10, e12)) = op2(h2(e10), e23)) | ~ (h12(op1(e10, e13)) = op2(e22, h2(e10))) | ~ (h12(op1(e11, e10)) = op2(h12(e11), h2(e10))) | ~ (h12(op1(e11, e12)) = op2(e23, h9(e11))) | ~ (h12(op1(e11, e13)) = op2(e23, h2(e10))) | ~ (h12(op1(e12, e10)) = op2(e23, h2(e10))) | ~ (h12(op1(e12, e11)) = op2(e23, h12(e11))) | ~ (h3(e10) = h12(op1(e12, e12))) | ~ (h12(e11) = h12(op1(e12, e13))) | ~ (h12(op1(e13, e10)) = op2(e22, h2(e10))) | ~ (h12(op1(e13, e11)) = op2(e22, h12(e11))) | ~ (h9(e11) = h12(op1(e13, e12))) | ~ (h2(e10) = h12(op1(e13, e13))) | sP40 | sP39 | ~ (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) | ~ (h12(op1(e10, e11)) = op2(h12(e10), h12(e11))) | ~ (h12(op1(e10, e10)) = op2(h12(e10), h12(e10)))), inference(forward_demodulation, [], [f2088, f1767])).
fof(f2088, plain, (~ (h12(op1(e10, e12)) = op2(h12(e10), e23)) | ~ (h12(op1(e10, e13)) = op2(e22, h2(e10))) | ~ (h12(op1(e11, e10)) = op2(h12(e11), h2(e10))) | ~ (h12(op1(e11, e12)) = op2(e23, h9(e11))) | ~ (h12(op1(e11, e13)) = op2(e23, h2(e10))) | ~ (h12(op1(e12, e10)) = op2(e23, h2(e10))) | ~ (h12(op1(e12, e11)) = op2(e23, h12(e11))) | ~ (h3(e10) = h12(op1(e12, e12))) | ~ (h12(e11) = h12(op1(e12, e13))) | ~ (h12(op1(e13, e10)) = op2(e22, h2(e10))) | ~ (h12(op1(e13, e11)) = op2(e22, h12(e11))) | ~ (h9(e11) = h12(op1(e13, e12))) | ~ (h2(e10) = h12(op1(e13, e13))) | sP40 | sP39 | ~ (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) | ~ (h12(op1(e10, e11)) = op2(h12(e10), h12(e11))) | ~ (h12(op1(e10, e10)) = op2(h12(e10), h12(e10)))), inference(forward_demodulation, [], [f2087, f590])).
fof(f2087, plain, (~ (h12(op1(e10, e13)) = op2(e22, h2(e10))) | ~ (h12(op1(e11, e10)) = op2(h12(e11), h2(e10))) | ~ (h12(op1(e11, e12)) = op2(e23, h9(e11))) | ~ (h12(op1(e11, e13)) = op2(e23, h2(e10))) | ~ (h12(op1(e12, e10)) = op2(e23, h2(e10))) | ~ (h12(op1(e12, e11)) = op2(e23, h12(e11))) | ~ (h3(e10) = h12(op1(e12, e12))) | ~ (h12(e11) = h12(op1(e12, e13))) | ~ (h12(op1(e13, e10)) = op2(e22, h2(e10))) | ~ (h12(op1(e13, e11)) = op2(e22, h12(e11))) | ~ (h9(e11) = h12(op1(e13, e12))) | ~ (h2(e10) = h12(op1(e13, e13))) | sP40 | sP39 | ~ (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) | ~ (h12(op1(e10, e12)) = op2(h12(e10), h12(e12))) | ~ (h12(op1(e10, e11)) = op2(h12(e10), h12(e11))) | ~ (h12(op1(e10, e10)) = op2(h12(e10), h12(e10)))), inference(forward_demodulation, [], [f2086, f1627])).
fof(f1627, plain, (op2(h2(e10), e22) = op2(e22, h2(e10))), inference(backward_demodulation, [], [f289, f552])).
fof(f289, plain, (op2(op2(e22, e22), e22) = op2(e22, op2(e22, e22))), inference(cnf_transformation, [], [f6])).
fof(f2086, plain, (~ (h12(op1(e10, e13)) = op2(h2(e10), e22)) | ~ (h12(op1(e11, e10)) = op2(h12(e11), h2(e10))) | ~ (h12(op1(e11, e12)) = op2(e23, h9(e11))) | ~ (h12(op1(e11, e13)) = op2(e23, h2(e10))) | ~ (h12(op1(e12, e10)) = op2(e23, h2(e10))) | ~ (h12(op1(e12, e11)) = op2(e23, h12(e11))) | ~ (h3(e10) = h12(op1(e12, e12))) | ~ (h12(e11) = h12(op1(e12, e13))) | ~ (h12(op1(e13, e10)) = op2(e22, h2(e10))) | ~ (h12(op1(e13, e11)) = op2(e22, h12(e11))) | ~ (h9(e11) = h12(op1(e13, e12))) | ~ (h2(e10) = h12(op1(e13, e13))) | sP40 | sP39 | ~ (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) | ~ (h12(op1(e10, e12)) = op2(h12(e10), h12(e12))) | ~ (h12(op1(e10, e11)) = op2(h12(e10), h12(e11))) | ~ (h12(op1(e10, e10)) = op2(h12(e10), h12(e10)))), inference(forward_demodulation, [], [f2085, f1767])).
fof(f2085, plain, (~ (h12(op1(e10, e13)) = op2(h12(e10), e22)) | ~ (h12(op1(e11, e10)) = op2(h12(e11), h2(e10))) | ~ (h12(op1(e11, e12)) = op2(e23, h9(e11))) | ~ (h12(op1(e11, e13)) = op2(e23, h2(e10))) | ~ (h12(op1(e12, e10)) = op2(e23, h2(e10))) | ~ (h12(op1(e12, e11)) = op2(e23, h12(e11))) | ~ (h3(e10) = h12(op1(e12, e12))) | ~ (h12(e11) = h12(op1(e12, e13))) | ~ (h12(op1(e13, e10)) = op2(e22, h2(e10))) | ~ (h12(op1(e13, e11)) = op2(e22, h12(e11))) | ~ (h9(e11) = h12(op1(e13, e12))) | ~ (h2(e10) = h12(op1(e13, e13))) | sP40 | sP39 | ~ (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) | ~ (h12(op1(e10, e12)) = op2(h12(e10), h12(e12))) | ~ (h12(op1(e10, e11)) = op2(h12(e10), h12(e11))) | ~ (h12(op1(e10, e10)) = op2(h12(e10), h12(e10)))), inference(forward_demodulation, [], [f2084, f591])).
fof(f2084, plain, (~ (h12(op1(e11, e10)) = op2(h12(e11), h2(e10))) | ~ (h12(op1(e11, e12)) = op2(e23, h9(e11))) | ~ (h12(op1(e11, e13)) = op2(e23, h2(e10))) | ~ (h12(op1(e12, e10)) = op2(e23, h2(e10))) | ~ (h12(op1(e12, e11)) = op2(e23, h12(e11))) | ~ (h3(e10) = h12(op1(e12, e12))) | ~ (h12(e11) = h12(op1(e12, e13))) | ~ (h12(op1(e13, e10)) = op2(e22, h2(e10))) | ~ (h12(op1(e13, e11)) = op2(e22, h12(e11))) | ~ (h9(e11) = h12(op1(e13, e12))) | ~ (h2(e10) = h12(op1(e13, e13))) | sP40 | sP39 | ~ (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) | ~ (h12(op1(e10, e13)) = op2(h12(e10), h12(e13))) | ~ (h12(op1(e10, e12)) = op2(h12(e10), h12(e12))) | ~ (h12(op1(e10, e11)) = op2(h12(e10), h12(e11))) | ~ (h12(op1(e10, e10)) = op2(h12(e10), h12(e10)))), inference(forward_demodulation, [], [f2083, f1767])).
fof(f2083, plain, (~ (h12(op1(e11, e12)) = op2(e23, h9(e11))) | ~ (h12(op1(e11, e13)) = op2(e23, h2(e10))) | ~ (h12(op1(e12, e10)) = op2(e23, h2(e10))) | ~ (h12(op1(e12, e11)) = op2(e23, h12(e11))) | ~ (h3(e10) = h12(op1(e12, e12))) | ~ (h12(e11) = h12(op1(e12, e13))) | ~ (h12(op1(e13, e10)) = op2(e22, h2(e10))) | ~ (h12(op1(e13, e11)) = op2(e22, h12(e11))) | ~ (h9(e11) = h12(op1(e13, e12))) | ~ (h2(e10) = h12(op1(e13, e13))) | sP40 | sP39 | ~ (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) | ~ (h12(op1(e11, e10)) = op2(h12(e11), h12(e10))) | ~ (h12(op1(e10, e13)) = op2(h12(e10), h12(e13))) | ~ (h12(op1(e10, e12)) = op2(h12(e10), h12(e12))) | ~ (h12(op1(e10, e11)) = op2(h12(e10), h12(e11))) | ~ (h12(op1(e10, e10)) = op2(h12(e10), h12(e10)))), inference(forward_demodulation, [], [f2082, f1766])).
fof(f2082, plain, (~ (h12(op1(e11, e12)) = op2(h12(e11), e23)) | ~ (h12(op1(e11, e13)) = op2(e23, h2(e10))) | ~ (h12(op1(e12, e10)) = op2(e23, h2(e10))) | ~ (h12(op1(e12, e11)) = op2(e23, h12(e11))) | ~ (h3(e10) = h12(op1(e12, e12))) | ~ (h12(e11) = h12(op1(e12, e13))) | ~ (h12(op1(e13, e10)) = op2(e22, h2(e10))) | ~ (h12(op1(e13, e11)) = op2(e22, h12(e11))) | ~ (h9(e11) = h12(op1(e13, e12))) | ~ (h2(e10) = h12(op1(e13, e13))) | sP40 | sP39 | ~ (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) | ~ (h12(op1(e11, e10)) = op2(h12(e11), h12(e10))) | ~ (h12(op1(e10, e13)) = op2(h12(e10), h12(e13))) | ~ (h12(op1(e10, e12)) = op2(h12(e10), h12(e12))) | ~ (h12(op1(e10, e11)) = op2(h12(e10), h12(e11))) | ~ (h12(op1(e10, e10)) = op2(h12(e10), h12(e10)))), inference(forward_demodulation, [], [f2081, f590])).
fof(f2081, plain, (~ (h12(op1(e11, e13)) = op2(e23, h2(e10))) | ~ (h12(op1(e12, e10)) = op2(e23, h2(e10))) | ~ (h12(op1(e12, e11)) = op2(e23, h12(e11))) | ~ (h3(e10) = h12(op1(e12, e12))) | ~ (h12(e11) = h12(op1(e12, e13))) | ~ (h12(op1(e13, e10)) = op2(e22, h2(e10))) | ~ (h12(op1(e13, e11)) = op2(e22, h12(e11))) | ~ (h9(e11) = h12(op1(e13, e12))) | ~ (h2(e10) = h12(op1(e13, e13))) | sP40 | sP39 | ~ (h12(op1(e11, e12)) = op2(h12(e11), h12(e12))) | ~ (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) | ~ (h12(op1(e11, e10)) = op2(h12(e11), h12(e10))) | ~ (h12(op1(e10, e13)) = op2(h12(e10), h12(e13))) | ~ (h12(op1(e10, e12)) = op2(h12(e10), h12(e12))) | ~ (h12(op1(e10, e11)) = op2(h12(e10), h12(e11))) | ~ (h12(op1(e10, e10)) = op2(h12(e10), h12(e10)))), inference(forward_demodulation, [], [f2080, f1759])).
fof(f1759, plain, (op2(e23, h2(e10)) = op2(h12(e11), e22)), inference(backward_demodulation, [], [f1629, f593])).
fof(f1629, plain, (op2(op2(e23, e22), e22) = op2(e23, h2(e10))), inference(backward_demodulation, [], [f305, f552])).
fof(f305, plain, (op2(op2(e23, e22), e22) = op2(e23, op2(e22, e22))), inference(cnf_transformation, [], [f6])).
fof(f2080, plain, (~ (h12(op1(e11, e13)) = op2(h12(e11), e22)) | ~ (h12(op1(e12, e10)) = op2(e23, h2(e10))) | ~ (h12(op1(e12, e11)) = op2(e23, h12(e11))) | ~ (h3(e10) = h12(op1(e12, e12))) | ~ (h12(e11) = h12(op1(e12, e13))) | ~ (h12(op1(e13, e10)) = op2(e22, h2(e10))) | ~ (h12(op1(e13, e11)) = op2(e22, h12(e11))) | ~ (h9(e11) = h12(op1(e13, e12))) | ~ (h2(e10) = h12(op1(e13, e13))) | sP40 | sP39 | ~ (h12(op1(e11, e12)) = op2(h12(e11), h12(e12))) | ~ (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) | ~ (h12(op1(e11, e10)) = op2(h12(e11), h12(e10))) | ~ (h12(op1(e10, e13)) = op2(h12(e10), h12(e13))) | ~ (h12(op1(e10, e12)) = op2(h12(e10), h12(e12))) | ~ (h12(op1(e10, e11)) = op2(h12(e10), h12(e11))) | ~ (h12(op1(e10, e10)) = op2(h12(e10), h12(e10)))), inference(forward_demodulation, [], [f2079, f591])).
fof(f2079, plain, (~ (h12(op1(e12, e10)) = op2(e23, h2(e10))) | ~ (h12(op1(e12, e11)) = op2(e23, h12(e11))) | ~ (h3(e10) = h12(op1(e12, e12))) | ~ (h12(e11) = h12(op1(e12, e13))) | ~ (h12(op1(e13, e10)) = op2(e22, h2(e10))) | ~ (h12(op1(e13, e11)) = op2(e22, h12(e11))) | ~ (h9(e11) = h12(op1(e13, e12))) | ~ (h2(e10) = h12(op1(e13, e13))) | sP40 | sP39 | ~ (h12(op1(e11, e13)) = op2(h12(e11), h12(e13))) | ~ (h12(op1(e11, e12)) = op2(h12(e11), h12(e12))) | ~ (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) | ~ (h12(op1(e11, e10)) = op2(h12(e11), h12(e10))) | ~ (h12(op1(e10, e13)) = op2(h12(e10), h12(e13))) | ~ (h12(op1(e10, e12)) = op2(h12(e10), h12(e12))) | ~ (h12(op1(e10, e11)) = op2(h12(e10), h12(e11))) | ~ (h12(op1(e10, e10)) = op2(h12(e10), h12(e10)))), inference(forward_demodulation, [], [f2078, f590])).
fof(f2078, plain, (~ (h12(op1(e12, e10)) = op2(h12(e12), h2(e10))) | ~ (h12(op1(e12, e11)) = op2(e23, h12(e11))) | ~ (h3(e10) = h12(op1(e12, e12))) | ~ (h12(e11) = h12(op1(e12, e13))) | ~ (h12(op1(e13, e10)) = op2(e22, h2(e10))) | ~ (h12(op1(e13, e11)) = op2(e22, h12(e11))) | ~ (h9(e11) = h12(op1(e13, e12))) | ~ (h2(e10) = h12(op1(e13, e13))) | sP40 | sP39 | ~ (h12(op1(e11, e13)) = op2(h12(e11), h12(e13))) | ~ (h12(op1(e11, e12)) = op2(h12(e11), h12(e12))) | ~ (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) | ~ (h12(op1(e11, e10)) = op2(h12(e11), h12(e10))) | ~ (h12(op1(e10, e13)) = op2(h12(e10), h12(e13))) | ~ (h12(op1(e10, e12)) = op2(h12(e10), h12(e12))) | ~ (h12(op1(e10, e11)) = op2(h12(e10), h12(e11))) | ~ (h12(op1(e10, e10)) = op2(h12(e10), h12(e10)))), inference(forward_demodulation, [], [f2077, f1767])).
fof(f2077, plain, (~ (h12(op1(e12, e11)) = op2(e23, h12(e11))) | ~ (h3(e10) = h12(op1(e12, e12))) | ~ (h12(e11) = h12(op1(e12, e13))) | ~ (h12(op1(e13, e10)) = op2(e22, h2(e10))) | ~ (h12(op1(e13, e11)) = op2(e22, h12(e11))) | ~ (h9(e11) = h12(op1(e13, e12))) | ~ (h2(e10) = h12(op1(e13, e13))) | sP40 | sP39 | ~ (h12(op1(e12, e10)) = op2(h12(e12), h12(e10))) | ~ (h12(op1(e11, e13)) = op2(h12(e11), h12(e13))) | ~ (h12(op1(e11, e12)) = op2(h12(e11), h12(e12))) | ~ (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) | ~ (h12(op1(e11, e10)) = op2(h12(e11), h12(e10))) | ~ (h12(op1(e10, e13)) = op2(h12(e10), h12(e13))) | ~ (h12(op1(e10, e12)) = op2(h12(e10), h12(e12))) | ~ (h12(op1(e10, e11)) = op2(h12(e10), h12(e11))) | ~ (h12(op1(e10, e10)) = op2(h12(e10), h12(e10)))), inference(forward_demodulation, [], [f2076, f590])).
fof(f2076, plain, (~ (h3(e10) = h12(op1(e12, e12))) | ~ (h12(e11) = h12(op1(e12, e13))) | ~ (h12(op1(e13, e10)) = op2(e22, h2(e10))) | ~ (h12(op1(e13, e11)) = op2(e22, h12(e11))) | ~ (h9(e11) = h12(op1(e13, e12))) | ~ (h2(e10) = h12(op1(e13, e13))) | sP40 | sP39 | ~ (h12(op1(e12, e11)) = op2(h12(e12), h12(e11))) | ~ (h12(op1(e12, e10)) = op2(h12(e12), h12(e10))) | ~ (h12(op1(e11, e13)) = op2(h12(e11), h12(e13))) | ~ (h12(op1(e11, e12)) = op2(h12(e11), h12(e12))) | ~ (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) | ~ (h12(op1(e11, e10)) = op2(h12(e11), h12(e10))) | ~ (h12(op1(e10, e13)) = op2(h12(e10), h12(e13))) | ~ (h12(op1(e10, e12)) = op2(h12(e10), h12(e12))) | ~ (h12(op1(e10, e11)) = op2(h12(e10), h12(e11))) | ~ (h12(op1(e10, e10)) = op2(h12(e10), h12(e10)))), inference(forward_demodulation, [], [f2075, f556])).
fof(f2075, plain, (~ (op2(e23, e23) = h12(op1(e12, e12))) | ~ (h12(e11) = h12(op1(e12, e13))) | ~ (h12(op1(e13, e10)) = op2(e22, h2(e10))) | ~ (h12(op1(e13, e11)) = op2(e22, h12(e11))) | ~ (h9(e11) = h12(op1(e13, e12))) | ~ (h2(e10) = h12(op1(e13, e13))) | sP40 | sP39 | ~ (h12(op1(e12, e11)) = op2(h12(e12), h12(e11))) | ~ (h12(op1(e12, e10)) = op2(h12(e12), h12(e10))) | ~ (h12(op1(e11, e13)) = op2(h12(e11), h12(e13))) | ~ (h12(op1(e11, e12)) = op2(h12(e11), h12(e12))) | ~ (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) | ~ (h12(op1(e11, e10)) = op2(h12(e11), h12(e10))) | ~ (h12(op1(e10, e13)) = op2(h12(e10), h12(e13))) | ~ (h12(op1(e10, e12)) = op2(h12(e10), h12(e12))) | ~ (h12(op1(e10, e11)) = op2(h12(e10), h12(e11))) | ~ (h12(op1(e10, e10)) = op2(h12(e10), h12(e10)))), inference(forward_demodulation, [], [f2074, f590])).
fof(f2074, plain, (~ (h12(e11) = h12(op1(e12, e13))) | ~ (h12(op1(e13, e10)) = op2(e22, h2(e10))) | ~ (h12(op1(e13, e11)) = op2(e22, h12(e11))) | ~ (h9(e11) = h12(op1(e13, e12))) | ~ (h2(e10) = h12(op1(e13, e13))) | sP40 | sP39 | ~ (h12(op1(e12, e12)) = op2(h12(e12), h12(e12))) | ~ (h12(op1(e12, e11)) = op2(h12(e12), h12(e11))) | ~ (h12(op1(e12, e10)) = op2(h12(e12), h12(e10))) | ~ (h12(op1(e11, e13)) = op2(h12(e11), h12(e13))) | ~ (h12(op1(e11, e12)) = op2(h12(e11), h12(e12))) | ~ (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) | ~ (h12(op1(e11, e10)) = op2(h12(e11), h12(e10))) | ~ (h12(op1(e10, e13)) = op2(h12(e10), h12(e13))) | ~ (h12(op1(e10, e12)) = op2(h12(e10), h12(e12))) | ~ (h12(op1(e10, e11)) = op2(h12(e10), h12(e11))) | ~ (h12(op1(e10, e10)) = op2(h12(e10), h12(e10)))), inference(forward_demodulation, [], [f2073, f593])).
fof(f2073, plain, (~ (op2(e23, e22) = h12(op1(e12, e13))) | ~ (h12(op1(e13, e10)) = op2(e22, h2(e10))) | ~ (h12(op1(e13, e11)) = op2(e22, h12(e11))) | ~ (h9(e11) = h12(op1(e13, e12))) | ~ (h2(e10) = h12(op1(e13, e13))) | sP40 | sP39 | ~ (h12(op1(e12, e12)) = op2(h12(e12), h12(e12))) | ~ (h12(op1(e12, e11)) = op2(h12(e12), h12(e11))) | ~ (h12(op1(e12, e10)) = op2(h12(e12), h12(e10))) | ~ (h12(op1(e11, e13)) = op2(h12(e11), h12(e13))) | ~ (h12(op1(e11, e12)) = op2(h12(e11), h12(e12))) | ~ (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) | ~ (h12(op1(e11, e10)) = op2(h12(e11), h12(e10))) | ~ (h12(op1(e10, e13)) = op2(h12(e10), h12(e13))) | ~ (h12(op1(e10, e12)) = op2(h12(e10), h12(e12))) | ~ (h12(op1(e10, e11)) = op2(h12(e10), h12(e11))) | ~ (h12(op1(e10, e10)) = op2(h12(e10), h12(e10)))), inference(forward_demodulation, [], [f2072, f590])).
fof(f2072, plain, (~ (h12(op1(e12, e13)) = op2(h12(e12), e22)) | ~ (h12(op1(e13, e10)) = op2(e22, h2(e10))) | ~ (h12(op1(e13, e11)) = op2(e22, h12(e11))) | ~ (h9(e11) = h12(op1(e13, e12))) | ~ (h2(e10) = h12(op1(e13, e13))) | sP40 | sP39 | ~ (h12(op1(e12, e12)) = op2(h12(e12), h12(e12))) | ~ (h12(op1(e12, e11)) = op2(h12(e12), h12(e11))) | ~ (h12(op1(e12, e10)) = op2(h12(e12), h12(e10))) | ~ (h12(op1(e11, e13)) = op2(h12(e11), h12(e13))) | ~ (h12(op1(e11, e12)) = op2(h12(e11), h12(e12))) | ~ (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) | ~ (h12(op1(e11, e10)) = op2(h12(e11), h12(e10))) | ~ (h12(op1(e10, e13)) = op2(h12(e10), h12(e13))) | ~ (h12(op1(e10, e12)) = op2(h12(e10), h12(e12))) | ~ (h12(op1(e10, e11)) = op2(h12(e10), h12(e11))) | ~ (h12(op1(e10, e10)) = op2(h12(e10), h12(e10)))), inference(forward_demodulation, [], [f2071, f591])).
fof(f2071, plain, (~ (h12(op1(e13, e10)) = op2(e22, h2(e10))) | ~ (h12(op1(e13, e11)) = op2(e22, h12(e11))) | ~ (h9(e11) = h12(op1(e13, e12))) | ~ (h2(e10) = h12(op1(e13, e13))) | sP40 | sP39 | ~ (h12(op1(e12, e13)) = op2(h12(e12), h12(e13))) | ~ (h12(op1(e12, e12)) = op2(h12(e12), h12(e12))) | ~ (h12(op1(e12, e11)) = op2(h12(e12), h12(e11))) | ~ (h12(op1(e12, e10)) = op2(h12(e12), h12(e10))) | ~ (h12(op1(e11, e13)) = op2(h12(e11), h12(e13))) | ~ (h12(op1(e11, e12)) = op2(h12(e11), h12(e12))) | ~ (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) | ~ (h12(op1(e11, e10)) = op2(h12(e11), h12(e10))) | ~ (h12(op1(e10, e13)) = op2(h12(e10), h12(e13))) | ~ (h12(op1(e10, e12)) = op2(h12(e10), h12(e12))) | ~ (h12(op1(e10, e11)) = op2(h12(e10), h12(e11))) | ~ (h12(op1(e10, e10)) = op2(h12(e10), h12(e10)))), inference(forward_demodulation, [], [f2070, f591])).
fof(f2070, plain, (~ (h12(op1(e13, e10)) = op2(h12(e13), h2(e10))) | ~ (h12(op1(e13, e11)) = op2(e22, h12(e11))) | ~ (h9(e11) = h12(op1(e13, e12))) | ~ (h2(e10) = h12(op1(e13, e13))) | sP40 | sP39 | ~ (h12(op1(e12, e13)) = op2(h12(e12), h12(e13))) | ~ (h12(op1(e12, e12)) = op2(h12(e12), h12(e12))) | ~ (h12(op1(e12, e11)) = op2(h12(e12), h12(e11))) | ~ (h12(op1(e12, e10)) = op2(h12(e12), h12(e10))) | ~ (h12(op1(e11, e13)) = op2(h12(e11), h12(e13))) | ~ (h12(op1(e11, e12)) = op2(h12(e11), h12(e12))) | ~ (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) | ~ (h12(op1(e11, e10)) = op2(h12(e11), h12(e10))) | ~ (h12(op1(e10, e13)) = op2(h12(e10), h12(e13))) | ~ (h12(op1(e10, e12)) = op2(h12(e10), h12(e12))) | ~ (h12(op1(e10, e11)) = op2(h12(e10), h12(e11))) | ~ (h12(op1(e10, e10)) = op2(h12(e10), h12(e10)))), inference(forward_demodulation, [], [f2069, f1767])).
fof(f2069, plain, (~ (h12(op1(e13, e11)) = op2(e22, h12(e11))) | ~ (h9(e11) = h12(op1(e13, e12))) | ~ (h2(e10) = h12(op1(e13, e13))) | sP40 | sP39 | ~ (h12(op1(e13, e10)) = op2(h12(e13), h12(e10))) | ~ (h12(op1(e12, e13)) = op2(h12(e12), h12(e13))) | ~ (h12(op1(e12, e12)) = op2(h12(e12), h12(e12))) | ~ (h12(op1(e12, e11)) = op2(h12(e12), h12(e11))) | ~ (h12(op1(e12, e10)) = op2(h12(e12), h12(e10))) | ~ (h12(op1(e11, e13)) = op2(h12(e11), h12(e13))) | ~ (h12(op1(e11, e12)) = op2(h12(e11), h12(e12))) | ~ (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) | ~ (h12(op1(e11, e10)) = op2(h12(e11), h12(e10))) | ~ (h12(op1(e10, e13)) = op2(h12(e10), h12(e13))) | ~ (h12(op1(e10, e12)) = op2(h12(e10), h12(e12))) | ~ (h12(op1(e10, e11)) = op2(h12(e10), h12(e11))) | ~ (h12(op1(e10, e10)) = op2(h12(e10), h12(e10)))), inference(forward_demodulation, [], [f2068, f591])).
fof(f2068, plain, (~ (h9(e11) = h12(op1(e13, e12))) | ~ (h2(e10) = h12(op1(e13, e13))) | sP40 | sP39 | ~ (h12(op1(e13, e11)) = op2(h12(e13), h12(e11))) | ~ (h12(op1(e13, e10)) = op2(h12(e13), h12(e10))) | ~ (h12(op1(e12, e13)) = op2(h12(e12), h12(e13))) | ~ (h12(op1(e12, e12)) = op2(h12(e12), h12(e12))) | ~ (h12(op1(e12, e11)) = op2(h12(e12), h12(e11))) | ~ (h12(op1(e12, e10)) = op2(h12(e12), h12(e10))) | ~ (h12(op1(e11, e13)) = op2(h12(e11), h12(e13))) | ~ (h12(op1(e11, e12)) = op2(h12(e11), h12(e12))) | ~ (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) | ~ (h12(op1(e11, e10)) = op2(h12(e11), h12(e10))) | ~ (h12(op1(e10, e13)) = op2(h12(e10), h12(e13))) | ~ (h12(op1(e10, e12)) = op2(h12(e10), h12(e12))) | ~ (h12(op1(e10, e11)) = op2(h12(e10), h12(e11))) | ~ (h12(op1(e10, e10)) = op2(h12(e10), h12(e10)))), inference(forward_demodulation, [], [f2067, f581])).
fof(f2067, plain, (~ (op2(e22, e23) = h12(op1(e13, e12))) | ~ (h2(e10) = h12(op1(e13, e13))) | sP40 | sP39 | ~ (h12(op1(e13, e11)) = op2(h12(e13), h12(e11))) | ~ (h12(op1(e13, e10)) = op2(h12(e13), h12(e10))) | ~ (h12(op1(e12, e13)) = op2(h12(e12), h12(e13))) | ~ (h12(op1(e12, e12)) = op2(h12(e12), h12(e12))) | ~ (h12(op1(e12, e11)) = op2(h12(e12), h12(e11))) | ~ (h12(op1(e12, e10)) = op2(h12(e12), h12(e10))) | ~ (h12(op1(e11, e13)) = op2(h12(e11), h12(e13))) | ~ (h12(op1(e11, e12)) = op2(h12(e11), h12(e12))) | ~ (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) | ~ (h12(op1(e11, e10)) = op2(h12(e11), h12(e10))) | ~ (h12(op1(e10, e13)) = op2(h12(e10), h12(e13))) | ~ (h12(op1(e10, e12)) = op2(h12(e10), h12(e12))) | ~ (h12(op1(e10, e11)) = op2(h12(e10), h12(e11))) | ~ (h12(op1(e10, e10)) = op2(h12(e10), h12(e10)))), inference(forward_demodulation, [], [f2066, f591])).
fof(f2066, plain, (~ (h12(op1(e13, e12)) = op2(h12(e13), e23)) | ~ (h2(e10) = h12(op1(e13, e13))) | sP40 | sP39 | ~ (h12(op1(e13, e11)) = op2(h12(e13), h12(e11))) | ~ (h12(op1(e13, e10)) = op2(h12(e13), h12(e10))) | ~ (h12(op1(e12, e13)) = op2(h12(e12), h12(e13))) | ~ (h12(op1(e12, e12)) = op2(h12(e12), h12(e12))) | ~ (h12(op1(e12, e11)) = op2(h12(e12), h12(e11))) | ~ (h12(op1(e12, e10)) = op2(h12(e12), h12(e10))) | ~ (h12(op1(e11, e13)) = op2(h12(e11), h12(e13))) | ~ (h12(op1(e11, e12)) = op2(h12(e11), h12(e12))) | ~ (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) | ~ (h12(op1(e11, e10)) = op2(h12(e11), h12(e10))) | ~ (h12(op1(e10, e13)) = op2(h12(e10), h12(e13))) | ~ (h12(op1(e10, e12)) = op2(h12(e10), h12(e12))) | ~ (h12(op1(e10, e11)) = op2(h12(e10), h12(e11))) | ~ (h12(op1(e10, e10)) = op2(h12(e10), h12(e10)))), inference(forward_demodulation, [], [f2065, f590])).
fof(f2065, plain, (~ (h2(e10) = h12(op1(e13, e13))) | sP40 | sP39 | ~ (h12(op1(e13, e12)) = op2(h12(e13), h12(e12))) | ~ (h12(op1(e13, e11)) = op2(h12(e13), h12(e11))) | ~ (h12(op1(e13, e10)) = op2(h12(e13), h12(e10))) | ~ (h12(op1(e12, e13)) = op2(h12(e12), h12(e13))) | ~ (h12(op1(e12, e12)) = op2(h12(e12), h12(e12))) | ~ (h12(op1(e12, e11)) = op2(h12(e12), h12(e11))) | ~ (h12(op1(e12, e10)) = op2(h12(e12), h12(e10))) | ~ (h12(op1(e11, e13)) = op2(h12(e11), h12(e13))) | ~ (h12(op1(e11, e12)) = op2(h12(e11), h12(e12))) | ~ (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) | ~ (h12(op1(e11, e10)) = op2(h12(e11), h12(e10))) | ~ (h12(op1(e10, e13)) = op2(h12(e10), h12(e13))) | ~ (h12(op1(e10, e12)) = op2(h12(e10), h12(e12))) | ~ (h12(op1(e10, e11)) = op2(h12(e10), h12(e11))) | ~ (h12(op1(e10, e10)) = op2(h12(e10), h12(e10)))), inference(forward_demodulation, [], [f2064, f552])).
fof(f2064, plain, (~ (op2(e22, e22) = h12(op1(e13, e13))) | sP40 | sP39 | ~ (h12(op1(e13, e12)) = op2(h12(e13), h12(e12))) | ~ (h12(op1(e13, e11)) = op2(h12(e13), h12(e11))) | ~ (h12(op1(e13, e10)) = op2(h12(e13), h12(e10))) | ~ (h12(op1(e12, e13)) = op2(h12(e12), h12(e13))) | ~ (h12(op1(e12, e12)) = op2(h12(e12), h12(e12))) | ~ (h12(op1(e12, e11)) = op2(h12(e12), h12(e11))) | ~ (h12(op1(e12, e10)) = op2(h12(e12), h12(e10))) | ~ (h12(op1(e11, e13)) = op2(h12(e11), h12(e13))) | ~ (h12(op1(e11, e12)) = op2(h12(e11), h12(e12))) | ~ (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) | ~ (h12(op1(e11, e10)) = op2(h12(e11), h12(e10))) | ~ (h12(op1(e10, e13)) = op2(h12(e10), h12(e13))) | ~ (h12(op1(e10, e12)) = op2(h12(e10), h12(e12))) | ~ (h12(op1(e10, e11)) = op2(h12(e10), h12(e11))) | ~ (h12(op1(e10, e10)) = op2(h12(e10), h12(e10)))), inference(forward_demodulation, [], [f2063, f591])).
fof(f2063, plain, (sP40 | sP39 | ~ (h12(op1(e13, e13)) = op2(h12(e13), h12(e13))) | ~ (h12(op1(e13, e12)) = op2(h12(e13), h12(e12))) | ~ (h12(op1(e13, e11)) = op2(h12(e13), h12(e11))) | ~ (h12(op1(e13, e10)) = op2(h12(e13), h12(e10))) | ~ (h12(op1(e12, e13)) = op2(h12(e12), h12(e13))) | ~ (h12(op1(e12, e12)) = op2(h12(e12), h12(e12))) | ~ (h12(op1(e12, e11)) = op2(h12(e12), h12(e11))) | ~ (h12(op1(e12, e10)) = op2(h12(e12), h12(e10))) | ~ (h12(op1(e11, e13)) = op2(h12(e11), h12(e13))) | ~ (h12(op1(e11, e12)) = op2(h12(e11), h12(e12))) | ~ (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) | ~ (h12(op1(e11, e10)) = op2(h12(e11), h12(e10))) | ~ (h12(op1(e10, e13)) = op2(h12(e10), h12(e13))) | ~ (h12(op1(e10, e12)) = op2(h12(e10), h12(e12))) | ~ (h12(op1(e10, e11)) = op2(h12(e10), h12(e11))) | ~ (h12(op1(e10, e10)) = op2(h12(e10), h12(e10)))), inference(subsumption_resolution, [], [f2062, f1768])).
fof(f1768, plain, ~ sP41, inference(subsumption_resolution, [], [f597, f591])).
fof(f597, plain, (~ (e22 = h12(e13)) | ~ sP41), inference(cnf_transformation, [], [f94])).
fof(f94, plain, ((~ (e22 = h12(e13)) & ~ (e22 = h12(e12)) & ~ (e22 = h12(e11)) & ~ (e22 = h12(e10))) | ~ sP41), inference(nnf_transformation, [], [f86])).
fof(f86, plain, ((~ (e22 = h12(e13)) & ~ (e22 = h12(e12)) & ~ (e22 = h12(e11)) & ~ (e22 = h12(e10))) | ~ sP41), inference(usedef, [], [e86])).
fof(e86, plain, (sP41 <=> (~ (e22 = h12(e13)) & ~ (e22 = h12(e12)) & ~ (e22 = h12(e11)) & ~ (e22 = h12(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP41])])).
fof(f2062, plain, (sP41 | sP40 | sP39 | ~ (h12(op1(e13, e13)) = op2(h12(e13), h12(e13))) | ~ (h12(op1(e13, e12)) = op2(h12(e13), h12(e12))) | ~ (h12(op1(e13, e11)) = op2(h12(e13), h12(e11))) | ~ (h12(op1(e13, e10)) = op2(h12(e13), h12(e10))) | ~ (h12(op1(e12, e13)) = op2(h12(e12), h12(e13))) | ~ (h12(op1(e12, e12)) = op2(h12(e12), h12(e12))) | ~ (h12(op1(e12, e11)) = op2(h12(e12), h12(e11))) | ~ (h12(op1(e12, e10)) = op2(h12(e12), h12(e10))) | ~ (h12(op1(e11, e13)) = op2(h12(e11), h12(e13))) | ~ (h12(op1(e11, e12)) = op2(h12(e11), h12(e12))) | ~ (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) | ~ (h12(op1(e11, e10)) = op2(h12(e11), h12(e10))) | ~ (h12(op1(e10, e13)) = op2(h12(e10), h12(e13))) | ~ (h12(op1(e10, e12)) = op2(h12(e10), h12(e12))) | ~ (h12(op1(e10, e11)) = op2(h12(e10), h12(e11))) | ~ (h12(op1(e10, e10)) = op2(h12(e10), h12(e10)))), inference(subsumption_resolution, [], [f784, f590])).
fof(f784, plain, (~ (e23 = h12(e12)) | sP41 | sP40 | sP39 | ~ (h12(op1(e13, e13)) = op2(h12(e13), h12(e13))) | ~ (h12(op1(e13, e12)) = op2(h12(e13), h12(e12))) | ~ (h12(op1(e13, e11)) = op2(h12(e13), h12(e11))) | ~ (h12(op1(e13, e10)) = op2(h12(e13), h12(e10))) | ~ (h12(op1(e12, e13)) = op2(h12(e12), h12(e13))) | ~ (h12(op1(e12, e12)) = op2(h12(e12), h12(e12))) | ~ (h12(op1(e12, e11)) = op2(h12(e12), h12(e11))) | ~ (h12(op1(e12, e10)) = op2(h12(e12), h12(e10))) | ~ (h12(op1(e11, e13)) = op2(h12(e11), h12(e13))) | ~ (h12(op1(e11, e12)) = op2(h12(e11), h12(e12))) | ~ (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) | ~ (h12(op1(e11, e10)) = op2(h12(e11), h12(e10))) | ~ (h12(op1(e10, e13)) = op2(h12(e10), h12(e13))) | ~ (h12(op1(e10, e12)) = op2(h12(e10), h12(e12))) | ~ (h12(op1(e10, e11)) = op2(h12(e10), h12(e11))) | ~ (h12(op1(e10, e10)) = op2(h12(e10), h12(e10)))), inference(cnf_transformation, [], [f87])).
fof(f87, plain, (((~ (e23 = h12(e13)) & ~ (e23 = h12(e12)) & ~ (e23 = h12(e11)) & ~ (e23 = h12(e10))) | sP41 | sP40 | sP39 | ~ (h12(op1(e13, e13)) = op2(h12(e13), h12(e13))) | ~ (h12(op1(e13, e12)) = op2(h12(e13), h12(e12))) | ~ (h12(op1(e13, e11)) = op2(h12(e13), h12(e11))) | ~ (h12(op1(e13, e10)) = op2(h12(e13), h12(e10))) | ~ (h12(op1(e12, e13)) = op2(h12(e12), h12(e13))) | ~ (h12(op1(e12, e12)) = op2(h12(e12), h12(e12))) | ~ (h12(op1(e12, e11)) = op2(h12(e12), h12(e11))) | ~ (h12(op1(e12, e10)) = op2(h12(e12), h12(e10))) | ~ (h12(op1(e11, e13)) = op2(h12(e11), h12(e13))) | ~ (h12(op1(e11, e12)) = op2(h12(e11), h12(e12))) | ~ (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) | ~ (h12(op1(e11, e10)) = op2(h12(e11), h12(e10))) | ~ (h12(op1(e10, e13)) = op2(h12(e10), h12(e13))) | ~ (h12(op1(e10, e12)) = op2(h12(e10), h12(e12))) | ~ (h12(op1(e10, e11)) = op2(h12(e10), h12(e11))) | ~ (h12(op1(e10, e10)) = op2(h12(e10), h12(e10)))) & ((~ (e23 = h11(e13)) & ~ (e23 = h11(e12)) & ~ (e23 = h11(e11)) & ~ (e23 = h11(e10))) | sP38 | sP37 | sP36 | ~ (h11(op1(e13, e13)) = op2(h11(e13), h11(e13))) | ~ (h11(op1(e13, e12)) = op2(h11(e13), h11(e12))) | ~ (h11(op1(e13, e11)) = op2(h11(e13), h11(e11))) | ~ (h11(op1(e13, e10)) = op2(h11(e13), h11(e10))) | ~ (h11(op1(e12, e13)) = op2(h11(e12), h11(e13))) | ~ (h11(op1(e12, e12)) = op2(h11(e12), h11(e12))) | ~ (h11(op1(e12, e11)) = op2(h11(e12), h11(e11))) | ~ (h11(op1(e12, e10)) = op2(h11(e12), h11(e10))) | ~ (h11(op1(e11, e13)) = op2(h11(e11), h11(e13))) | ~ (h11(op1(e11, e12)) = op2(h11(e11), h11(e12))) | ~ (h11(op1(e11, e11)) = op2(h11(e11), h11(e11))) | ~ (h11(op1(e11, e10)) = op2(h11(e11), h11(e10))) | ~ (h11(op1(e10, e13)) = op2(h11(e10), h11(e13))) | ~ (h11(op1(e10, e12)) = op2(h11(e10), h11(e12))) | ~ (h11(op1(e10, e11)) = op2(h11(e10), h11(e11))) | ~ (h11(op1(e10, e10)) = op2(h11(e10), h11(e10)))) & ((~ (e23 = h10(e13)) & ~ (e23 = h10(e12)) & ~ (e23 = h10(e11)) & ~ (e23 = h10(e10))) | sP35 | sP34 | sP33 | ~ (h10(op1(e13, e13)) = op2(h10(e13), h10(e13))) | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) | ~ (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) | ~ (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) | ~ (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))) & ((~ (e23 = h9(e13)) & ~ (e23 = h9(e12)) & ~ (e23 = h9(e11)) & ~ (e23 = h9(e10))) | sP32 | sP31 | sP30 | ~ (h9(op1(e13, e13)) = op2(h9(e13), h9(e13))) | ~ (h9(op1(e13, e12)) = op2(h9(e13), h9(e12))) | ~ (h9(op1(e13, e11)) = op2(h9(e13), h9(e11))) | ~ (h9(op1(e13, e10)) = op2(h9(e13), h9(e10))) | ~ (h9(op1(e12, e13)) = op2(h9(e12), h9(e13))) | ~ (h9(op1(e12, e12)) = op2(h9(e12), h9(e12))) | ~ (h9(op1(e12, e11)) = op2(h9(e12), h9(e11))) | ~ (h9(op1(e12, e10)) = op2(h9(e12), h9(e10))) | ~ (h9(op1(e11, e13)) = op2(h9(e11), h9(e13))) | ~ (h9(op1(e11, e12)) = op2(h9(e11), h9(e12))) | ~ (h9(op1(e11, e11)) = op2(h9(e11), h9(e11))) | ~ (h9(op1(e11, e10)) = op2(h9(e11), h9(e10))) | ~ (h9(op1(e10, e13)) = op2(h9(e10), h9(e13))) | ~ (h9(op1(e10, e12)) = op2(h9(e10), h9(e12))) | ~ (h9(op1(e10, e11)) = op2(h9(e10), h9(e11))) | ~ (h9(op1(e10, e10)) = op2(h9(e10), h9(e10)))) & ((~ (e23 = h8(e13)) & ~ (e23 = h8(e12)) & ~ (e23 = h8(e11)) & ~ (e23 = h8(e10))) | sP29 | sP28 | sP27 | ~ (h8(op1(e13, e13)) = op2(h8(e13), h8(e13))) | ~ (h8(op1(e13, e12)) = op2(h8(e13), h8(e12))) | ~ (h8(op1(e13, e11)) = op2(h8(e13), h8(e11))) | ~ (h8(op1(e13, e10)) = op2(h8(e13), h8(e10))) | ~ (h8(op1(e12, e13)) = op2(h8(e12), h8(e13))) | ~ (h8(op1(e12, e12)) = op2(h8(e12), h8(e12))) | ~ (h8(op1(e12, e11)) = op2(h8(e12), h8(e11))) | ~ (h8(op1(e12, e10)) = op2(h8(e12), h8(e10))) | ~ (h8(op1(e11, e13)) = op2(h8(e11), h8(e13))) | ~ (h8(op1(e11, e12)) = op2(h8(e11), h8(e12))) | ~ (h8(op1(e11, e11)) = op2(h8(e11), h8(e11))) | ~ (h8(op1(e11, e10)) = op2(h8(e11), h8(e10))) | ~ (h8(op1(e10, e13)) = op2(h8(e10), h8(e13))) | ~ (h8(op1(e10, e12)) = op2(h8(e10), h8(e12))) | ~ (h8(op1(e10, e11)) = op2(h8(e10), h8(e11))) | ~ (h8(op1(e10, e10)) = op2(h8(e10), h8(e10)))) & ((~ (e23 = h7(e13)) & ~ (e23 = h7(e12)) & ~ (e23 = h7(e11)) & ~ (e23 = h7(e10))) | sP26 | sP25 | sP24 | ~ (h7(op1(e13, e13)) = op2(h7(e13), h7(e13))) | ~ (h7(op1(e13, e12)) = op2(h7(e13), h7(e12))) | ~ (h7(op1(e13, e11)) = op2(h7(e13), h7(e11))) | ~ (h7(op1(e13, e10)) = op2(h7(e13), h7(e10))) | ~ (h7(op1(e12, e13)) = op2(h7(e12), h7(e13))) | ~ (h7(op1(e12, e12)) = op2(h7(e12), h7(e12))) | ~ (h7(op1(e12, e11)) = op2(h7(e12), h7(e11))) | ~ (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) | ~ (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) | ~ (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) | ~ (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) | ~ (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) | ~ (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) | ~ (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) | ~ (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))) & ((~ (e23 = h6(e13)) & ~ (e23 = h6(e12)) & ~ (e23 = h6(e11)) & ~ (e23 = h6(e10))) | sP23 | sP22 | sP21 | ~ (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) | ~ (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))) & ((~ (e23 = h5(e13)) & ~ (e23 = h5(e12)) & ~ (e23 = h5(e11)) & ~ (e23 = h5(e10))) | sP20 | sP19 | sP18 | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))) & ((~ (e23 = h4(e13)) & ~ (e23 = h4(e12)) & ~ (e23 = h4(e11)) & ~ (e23 = h4(e10))) | sP17 | sP16 | sP15 | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) & ((~ (e23 = h3(e13)) & ~ (e23 = h3(e12)) & ~ (e23 = h3(e11)) & ~ (e23 = h3(e10))) | sP14 | sP13 | sP12 | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) & ((~ (e23 = h2(e13)) & ~ (e23 = h2(e12)) & ~ (e23 = h2(e11)) & ~ (e23 = h2(e10))) | sP11 | sP10 | sP9 | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) & ((~ (e23 = h1(e13)) & ~ (e23 = h1(e12)) & ~ (e23 = h1(e11)) & ~ (e23 = h1(e10))) | sP8 | sP7 | sP6 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(definition_folding, [], [f42, e86, e85, e84, e83, e82, e81, e80, e79, e78, e77, e76, e75, e74, e73, e72, e71, e70, e69, e68, e67, e66, e65, e64, e63, e62, e61, e60, e59, e58, e57, e56, e55, e54, e53, e52, e51])).
fof(f51, plain, ((~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10))) | ~ sP6), inference(usedef, [], [e51])).
fof(e51, plain, (sP6 <=> (~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f52, plain, ((~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10))) | ~ sP7), inference(usedef, [], [e52])).
fof(e52, plain, (sP7 <=> (~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f53, plain, ((~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10))) | ~ sP8), inference(usedef, [], [e53])).
fof(e53, plain, (sP8 <=> (~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f54, plain, ((~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10))) | ~ sP9), inference(usedef, [], [e54])).
fof(e54, plain, (sP9 <=> (~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f55, plain, ((~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10))) | ~ sP10), inference(usedef, [], [e55])).
fof(e55, plain, (sP10 <=> (~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f56, plain, ((~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10))) | ~ sP11), inference(usedef, [], [e56])).
fof(e56, plain, (sP11 <=> (~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f57, plain, ((~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10))) | ~ sP12), inference(usedef, [], [e57])).
fof(e57, plain, (sP12 <=> (~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f58, plain, ((~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10))) | ~ sP13), inference(usedef, [], [e58])).
fof(e58, plain, (sP13 <=> (~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f59, plain, ((~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10))) | ~ sP14), inference(usedef, [], [e59])).
fof(e59, plain, (sP14 <=> (~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f60, plain, ((~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10))) | ~ sP15), inference(usedef, [], [e60])).
fof(e60, plain, (sP15 <=> (~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f61, plain, ((~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10))) | ~ sP16), inference(usedef, [], [e61])).
fof(e61, plain, (sP16 <=> (~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f62, plain, ((~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10))) | ~ sP17), inference(usedef, [], [e62])).
fof(e62, plain, (sP17 <=> (~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f63, plain, ((~ (e20 = h5(e13)) & ~ (e20 = h5(e12)) & ~ (e20 = h5(e11)) & ~ (e20 = h5(e10))) | ~ sP18), inference(usedef, [], [e63])).
fof(e63, plain, (sP18 <=> (~ (e20 = h5(e13)) & ~ (e20 = h5(e12)) & ~ (e20 = h5(e11)) & ~ (e20 = h5(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f64, plain, ((~ (e21 = h5(e13)) & ~ (e21 = h5(e12)) & ~ (e21 = h5(e11)) & ~ (e21 = h5(e10))) | ~ sP19), inference(usedef, [], [e64])).
fof(e64, plain, (sP19 <=> (~ (e21 = h5(e13)) & ~ (e21 = h5(e12)) & ~ (e21 = h5(e11)) & ~ (e21 = h5(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP19])])).
fof(f65, plain, ((~ (e22 = h5(e13)) & ~ (e22 = h5(e12)) & ~ (e22 = h5(e11)) & ~ (e22 = h5(e10))) | ~ sP20), inference(usedef, [], [e65])).
fof(e65, plain, (sP20 <=> (~ (e22 = h5(e13)) & ~ (e22 = h5(e12)) & ~ (e22 = h5(e11)) & ~ (e22 = h5(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP20])])).
fof(f66, plain, ((~ (e20 = h6(e13)) & ~ (e20 = h6(e12)) & ~ (e20 = h6(e11)) & ~ (e20 = h6(e10))) | ~ sP21), inference(usedef, [], [e66])).
fof(e66, plain, (sP21 <=> (~ (e20 = h6(e13)) & ~ (e20 = h6(e12)) & ~ (e20 = h6(e11)) & ~ (e20 = h6(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP21])])).
fof(f67, plain, ((~ (e21 = h6(e13)) & ~ (e21 = h6(e12)) & ~ (e21 = h6(e11)) & ~ (e21 = h6(e10))) | ~ sP22), inference(usedef, [], [e67])).
fof(e67, plain, (sP22 <=> (~ (e21 = h6(e13)) & ~ (e21 = h6(e12)) & ~ (e21 = h6(e11)) & ~ (e21 = h6(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP22])])).
fof(f68, plain, ((~ (e22 = h6(e13)) & ~ (e22 = h6(e12)) & ~ (e22 = h6(e11)) & ~ (e22 = h6(e10))) | ~ sP23), inference(usedef, [], [e68])).
fof(e68, plain, (sP23 <=> (~ (e22 = h6(e13)) & ~ (e22 = h6(e12)) & ~ (e22 = h6(e11)) & ~ (e22 = h6(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP23])])).
fof(f69, plain, ((~ (e20 = h7(e13)) & ~ (e20 = h7(e12)) & ~ (e20 = h7(e11)) & ~ (e20 = h7(e10))) | ~ sP24), inference(usedef, [], [e69])).
fof(e69, plain, (sP24 <=> (~ (e20 = h7(e13)) & ~ (e20 = h7(e12)) & ~ (e20 = h7(e11)) & ~ (e20 = h7(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP24])])).
fof(f70, plain, ((~ (e21 = h7(e13)) & ~ (e21 = h7(e12)) & ~ (e21 = h7(e11)) & ~ (e21 = h7(e10))) | ~ sP25), inference(usedef, [], [e70])).
fof(e70, plain, (sP25 <=> (~ (e21 = h7(e13)) & ~ (e21 = h7(e12)) & ~ (e21 = h7(e11)) & ~ (e21 = h7(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP25])])).
fof(f71, plain, ((~ (e22 = h7(e13)) & ~ (e22 = h7(e12)) & ~ (e22 = h7(e11)) & ~ (e22 = h7(e10))) | ~ sP26), inference(usedef, [], [e71])).
fof(e71, plain, (sP26 <=> (~ (e22 = h7(e13)) & ~ (e22 = h7(e12)) & ~ (e22 = h7(e11)) & ~ (e22 = h7(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP26])])).
fof(f72, plain, ((~ (e20 = h8(e13)) & ~ (e20 = h8(e12)) & ~ (e20 = h8(e11)) & ~ (e20 = h8(e10))) | ~ sP27), inference(usedef, [], [e72])).
fof(e72, plain, (sP27 <=> (~ (e20 = h8(e13)) & ~ (e20 = h8(e12)) & ~ (e20 = h8(e11)) & ~ (e20 = h8(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP27])])).
fof(f73, plain, ((~ (e21 = h8(e13)) & ~ (e21 = h8(e12)) & ~ (e21 = h8(e11)) & ~ (e21 = h8(e10))) | ~ sP28), inference(usedef, [], [e73])).
fof(e73, plain, (sP28 <=> (~ (e21 = h8(e13)) & ~ (e21 = h8(e12)) & ~ (e21 = h8(e11)) & ~ (e21 = h8(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP28])])).
fof(f74, plain, ((~ (e22 = h8(e13)) & ~ (e22 = h8(e12)) & ~ (e22 = h8(e11)) & ~ (e22 = h8(e10))) | ~ sP29), inference(usedef, [], [e74])).
fof(e74, plain, (sP29 <=> (~ (e22 = h8(e13)) & ~ (e22 = h8(e12)) & ~ (e22 = h8(e11)) & ~ (e22 = h8(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP29])])).
fof(f75, plain, ((~ (e20 = h9(e13)) & ~ (e20 = h9(e12)) & ~ (e20 = h9(e11)) & ~ (e20 = h9(e10))) | ~ sP30), inference(usedef, [], [e75])).
fof(e75, plain, (sP30 <=> (~ (e20 = h9(e13)) & ~ (e20 = h9(e12)) & ~ (e20 = h9(e11)) & ~ (e20 = h9(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP30])])).
fof(f76, plain, ((~ (e21 = h9(e13)) & ~ (e21 = h9(e12)) & ~ (e21 = h9(e11)) & ~ (e21 = h9(e10))) | ~ sP31), inference(usedef, [], [e76])).
fof(e76, plain, (sP31 <=> (~ (e21 = h9(e13)) & ~ (e21 = h9(e12)) & ~ (e21 = h9(e11)) & ~ (e21 = h9(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP31])])).
fof(f77, plain, ((~ (e22 = h9(e13)) & ~ (e22 = h9(e12)) & ~ (e22 = h9(e11)) & ~ (e22 = h9(e10))) | ~ sP32), inference(usedef, [], [e77])).
fof(e77, plain, (sP32 <=> (~ (e22 = h9(e13)) & ~ (e22 = h9(e12)) & ~ (e22 = h9(e11)) & ~ (e22 = h9(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP32])])).
fof(f78, plain, ((~ (e20 = h10(e13)) & ~ (e20 = h10(e12)) & ~ (e20 = h10(e11)) & ~ (e20 = h10(e10))) | ~ sP33), inference(usedef, [], [e78])).
fof(e78, plain, (sP33 <=> (~ (e20 = h10(e13)) & ~ (e20 = h10(e12)) & ~ (e20 = h10(e11)) & ~ (e20 = h10(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP33])])).
fof(f79, plain, ((~ (e21 = h10(e13)) & ~ (e21 = h10(e12)) & ~ (e21 = h10(e11)) & ~ (e21 = h10(e10))) | ~ sP34), inference(usedef, [], [e79])).
fof(e79, plain, (sP34 <=> (~ (e21 = h10(e13)) & ~ (e21 = h10(e12)) & ~ (e21 = h10(e11)) & ~ (e21 = h10(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP34])])).
fof(f80, plain, ((~ (e22 = h10(e13)) & ~ (e22 = h10(e12)) & ~ (e22 = h10(e11)) & ~ (e22 = h10(e10))) | ~ sP35), inference(usedef, [], [e80])).
fof(e80, plain, (sP35 <=> (~ (e22 = h10(e13)) & ~ (e22 = h10(e12)) & ~ (e22 = h10(e11)) & ~ (e22 = h10(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP35])])).
fof(f81, plain, ((~ (e20 = h11(e13)) & ~ (e20 = h11(e12)) & ~ (e20 = h11(e11)) & ~ (e20 = h11(e10))) | ~ sP36), inference(usedef, [], [e81])).
fof(e81, plain, (sP36 <=> (~ (e20 = h11(e13)) & ~ (e20 = h11(e12)) & ~ (e20 = h11(e11)) & ~ (e20 = h11(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP36])])).
fof(f82, plain, ((~ (e21 = h11(e13)) & ~ (e21 = h11(e12)) & ~ (e21 = h11(e11)) & ~ (e21 = h11(e10))) | ~ sP37), inference(usedef, [], [e82])).
fof(e82, plain, (sP37 <=> (~ (e21 = h11(e13)) & ~ (e21 = h11(e12)) & ~ (e21 = h11(e11)) & ~ (e21 = h11(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP37])])).
fof(f83, plain, ((~ (e22 = h11(e13)) & ~ (e22 = h11(e12)) & ~ (e22 = h11(e11)) & ~ (e22 = h11(e10))) | ~ sP38), inference(usedef, [], [e83])).
fof(e83, plain, (sP38 <=> (~ (e22 = h11(e13)) & ~ (e22 = h11(e12)) & ~ (e22 = h11(e11)) & ~ (e22 = h11(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP38])])).
fof(f84, plain, ((~ (e20 = h12(e13)) & ~ (e20 = h12(e12)) & ~ (e20 = h12(e11)) & ~ (e20 = h12(e10))) | ~ sP39), inference(usedef, [], [e84])).
fof(e84, plain, (sP39 <=> (~ (e20 = h12(e13)) & ~ (e20 = h12(e12)) & ~ (e20 = h12(e11)) & ~ (e20 = h12(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP39])])).
fof(f85, plain, ((~ (e21 = h12(e13)) & ~ (e21 = h12(e12)) & ~ (e21 = h12(e11)) & ~ (e21 = h12(e10))) | ~ sP40), inference(usedef, [], [e85])).
fof(e85, plain, (sP40 <=> (~ (e21 = h12(e13)) & ~ (e21 = h12(e12)) & ~ (e21 = h12(e11)) & ~ (e21 = h12(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP40])])).
fof(f42, plain, (((~ (e23 = h12(e13)) & ~ (e23 = h12(e12)) & ~ (e23 = h12(e11)) & ~ (e23 = h12(e10))) | (~ (e22 = h12(e13)) & ~ (e22 = h12(e12)) & ~ (e22 = h12(e11)) & ~ (e22 = h12(e10))) | (~ (e21 = h12(e13)) & ~ (e21 = h12(e12)) & ~ (e21 = h12(e11)) & ~ (e21 = h12(e10))) | (~ (e20 = h12(e13)) & ~ (e20 = h12(e12)) & ~ (e20 = h12(e11)) & ~ (e20 = h12(e10))) | ~ (h12(op1(e13, e13)) = op2(h12(e13), h12(e13))) | ~ (h12(op1(e13, e12)) = op2(h12(e13), h12(e12))) | ~ (h12(op1(e13, e11)) = op2(h12(e13), h12(e11))) | ~ (h12(op1(e13, e10)) = op2(h12(e13), h12(e10))) | ~ (h12(op1(e12, e13)) = op2(h12(e12), h12(e13))) | ~ (h12(op1(e12, e12)) = op2(h12(e12), h12(e12))) | ~ (h12(op1(e12, e11)) = op2(h12(e12), h12(e11))) | ~ (h12(op1(e12, e10)) = op2(h12(e12), h12(e10))) | ~ (h12(op1(e11, e13)) = op2(h12(e11), h12(e13))) | ~ (h12(op1(e11, e12)) = op2(h12(e11), h12(e12))) | ~ (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) | ~ (h12(op1(e11, e10)) = op2(h12(e11), h12(e10))) | ~ (h12(op1(e10, e13)) = op2(h12(e10), h12(e13))) | ~ (h12(op1(e10, e12)) = op2(h12(e10), h12(e12))) | ~ (h12(op1(e10, e11)) = op2(h12(e10), h12(e11))) | ~ (h12(op1(e10, e10)) = op2(h12(e10), h12(e10)))) & ((~ (e23 = h11(e13)) & ~ (e23 = h11(e12)) & ~ (e23 = h11(e11)) & ~ (e23 = h11(e10))) | (~ (e22 = h11(e13)) & ~ (e22 = h11(e12)) & ~ (e22 = h11(e11)) & ~ (e22 = h11(e10))) | (~ (e21 = h11(e13)) & ~ (e21 = h11(e12)) & ~ (e21 = h11(e11)) & ~ (e21 = h11(e10))) | (~ (e20 = h11(e13)) & ~ (e20 = h11(e12)) & ~ (e20 = h11(e11)) & ~ (e20 = h11(e10))) | ~ (h11(op1(e13, e13)) = op2(h11(e13), h11(e13))) | ~ (h11(op1(e13, e12)) = op2(h11(e13), h11(e12))) | ~ (h11(op1(e13, e11)) = op2(h11(e13), h11(e11))) | ~ (h11(op1(e13, e10)) = op2(h11(e13), h11(e10))) | ~ (h11(op1(e12, e13)) = op2(h11(e12), h11(e13))) | ~ (h11(op1(e12, e12)) = op2(h11(e12), h11(e12))) | ~ (h11(op1(e12, e11)) = op2(h11(e12), h11(e11))) | ~ (h11(op1(e12, e10)) = op2(h11(e12), h11(e10))) | ~ (h11(op1(e11, e13)) = op2(h11(e11), h11(e13))) | ~ (h11(op1(e11, e12)) = op2(h11(e11), h11(e12))) | ~ (h11(op1(e11, e11)) = op2(h11(e11), h11(e11))) | ~ (h11(op1(e11, e10)) = op2(h11(e11), h11(e10))) | ~ (h11(op1(e10, e13)) = op2(h11(e10), h11(e13))) | ~ (h11(op1(e10, e12)) = op2(h11(e10), h11(e12))) | ~ (h11(op1(e10, e11)) = op2(h11(e10), h11(e11))) | ~ (h11(op1(e10, e10)) = op2(h11(e10), h11(e10)))) & ((~ (e23 = h10(e13)) & ~ (e23 = h10(e12)) & ~ (e23 = h10(e11)) & ~ (e23 = h10(e10))) | (~ (e22 = h10(e13)) & ~ (e22 = h10(e12)) & ~ (e22 = h10(e11)) & ~ (e22 = h10(e10))) | (~ (e21 = h10(e13)) & ~ (e21 = h10(e12)) & ~ (e21 = h10(e11)) & ~ (e21 = h10(e10))) | (~ (e20 = h10(e13)) & ~ (e20 = h10(e12)) & ~ (e20 = h10(e11)) & ~ (e20 = h10(e10))) | ~ (h10(op1(e13, e13)) = op2(h10(e13), h10(e13))) | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) | ~ (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) | ~ (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) | ~ (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))) & ((~ (e23 = h9(e13)) & ~ (e23 = h9(e12)) & ~ (e23 = h9(e11)) & ~ (e23 = h9(e10))) | (~ (e22 = h9(e13)) & ~ (e22 = h9(e12)) & ~ (e22 = h9(e11)) & ~ (e22 = h9(e10))) | (~ (e21 = h9(e13)) & ~ (e21 = h9(e12)) & ~ (e21 = h9(e11)) & ~ (e21 = h9(e10))) | (~ (e20 = h9(e13)) & ~ (e20 = h9(e12)) & ~ (e20 = h9(e11)) & ~ (e20 = h9(e10))) | ~ (h9(op1(e13, e13)) = op2(h9(e13), h9(e13))) | ~ (h9(op1(e13, e12)) = op2(h9(e13), h9(e12))) | ~ (h9(op1(e13, e11)) = op2(h9(e13), h9(e11))) | ~ (h9(op1(e13, e10)) = op2(h9(e13), h9(e10))) | ~ (h9(op1(e12, e13)) = op2(h9(e12), h9(e13))) | ~ (h9(op1(e12, e12)) = op2(h9(e12), h9(e12))) | ~ (h9(op1(e12, e11)) = op2(h9(e12), h9(e11))) | ~ (h9(op1(e12, e10)) = op2(h9(e12), h9(e10))) | ~ (h9(op1(e11, e13)) = op2(h9(e11), h9(e13))) | ~ (h9(op1(e11, e12)) = op2(h9(e11), h9(e12))) | ~ (h9(op1(e11, e11)) = op2(h9(e11), h9(e11))) | ~ (h9(op1(e11, e10)) = op2(h9(e11), h9(e10))) | ~ (h9(op1(e10, e13)) = op2(h9(e10), h9(e13))) | ~ (h9(op1(e10, e12)) = op2(h9(e10), h9(e12))) | ~ (h9(op1(e10, e11)) = op2(h9(e10), h9(e11))) | ~ (h9(op1(e10, e10)) = op2(h9(e10), h9(e10)))) & ((~ (e23 = h8(e13)) & ~ (e23 = h8(e12)) & ~ (e23 = h8(e11)) & ~ (e23 = h8(e10))) | (~ (e22 = h8(e13)) & ~ (e22 = h8(e12)) & ~ (e22 = h8(e11)) & ~ (e22 = h8(e10))) | (~ (e21 = h8(e13)) & ~ (e21 = h8(e12)) & ~ (e21 = h8(e11)) & ~ (e21 = h8(e10))) | (~ (e20 = h8(e13)) & ~ (e20 = h8(e12)) & ~ (e20 = h8(e11)) & ~ (e20 = h8(e10))) | ~ (h8(op1(e13, e13)) = op2(h8(e13), h8(e13))) | ~ (h8(op1(e13, e12)) = op2(h8(e13), h8(e12))) | ~ (h8(op1(e13, e11)) = op2(h8(e13), h8(e11))) | ~ (h8(op1(e13, e10)) = op2(h8(e13), h8(e10))) | ~ (h8(op1(e12, e13)) = op2(h8(e12), h8(e13))) | ~ (h8(op1(e12, e12)) = op2(h8(e12), h8(e12))) | ~ (h8(op1(e12, e11)) = op2(h8(e12), h8(e11))) | ~ (h8(op1(e12, e10)) = op2(h8(e12), h8(e10))) | ~ (h8(op1(e11, e13)) = op2(h8(e11), h8(e13))) | ~ (h8(op1(e11, e12)) = op2(h8(e11), h8(e12))) | ~ (h8(op1(e11, e11)) = op2(h8(e11), h8(e11))) | ~ (h8(op1(e11, e10)) = op2(h8(e11), h8(e10))) | ~ (h8(op1(e10, e13)) = op2(h8(e10), h8(e13))) | ~ (h8(op1(e10, e12)) = op2(h8(e10), h8(e12))) | ~ (h8(op1(e10, e11)) = op2(h8(e10), h8(e11))) | ~ (h8(op1(e10, e10)) = op2(h8(e10), h8(e10)))) & ((~ (e23 = h7(e13)) & ~ (e23 = h7(e12)) & ~ (e23 = h7(e11)) & ~ (e23 = h7(e10))) | (~ (e22 = h7(e13)) & ~ (e22 = h7(e12)) & ~ (e22 = h7(e11)) & ~ (e22 = h7(e10))) | (~ (e21 = h7(e13)) & ~ (e21 = h7(e12)) & ~ (e21 = h7(e11)) & ~ (e21 = h7(e10))) | (~ (e20 = h7(e13)) & ~ (e20 = h7(e12)) & ~ (e20 = h7(e11)) & ~ (e20 = h7(e10))) | ~ (h7(op1(e13, e13)) = op2(h7(e13), h7(e13))) | ~ (h7(op1(e13, e12)) = op2(h7(e13), h7(e12))) | ~ (h7(op1(e13, e11)) = op2(h7(e13), h7(e11))) | ~ (h7(op1(e13, e10)) = op2(h7(e13), h7(e10))) | ~ (h7(op1(e12, e13)) = op2(h7(e12), h7(e13))) | ~ (h7(op1(e12, e12)) = op2(h7(e12), h7(e12))) | ~ (h7(op1(e12, e11)) = op2(h7(e12), h7(e11))) | ~ (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) | ~ (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) | ~ (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) | ~ (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) | ~ (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) | ~ (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) | ~ (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) | ~ (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))) & ((~ (e23 = h6(e13)) & ~ (e23 = h6(e12)) & ~ (e23 = h6(e11)) & ~ (e23 = h6(e10))) | (~ (e22 = h6(e13)) & ~ (e22 = h6(e12)) & ~ (e22 = h6(e11)) & ~ (e22 = h6(e10))) | (~ (e21 = h6(e13)) & ~ (e21 = h6(e12)) & ~ (e21 = h6(e11)) & ~ (e21 = h6(e10))) | (~ (e20 = h6(e13)) & ~ (e20 = h6(e12)) & ~ (e20 = h6(e11)) & ~ (e20 = h6(e10))) | ~ (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) | ~ (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))) & ((~ (e23 = h5(e13)) & ~ (e23 = h5(e12)) & ~ (e23 = h5(e11)) & ~ (e23 = h5(e10))) | (~ (e22 = h5(e13)) & ~ (e22 = h5(e12)) & ~ (e22 = h5(e11)) & ~ (e22 = h5(e10))) | (~ (e21 = h5(e13)) & ~ (e21 = h5(e12)) & ~ (e21 = h5(e11)) & ~ (e21 = h5(e10))) | (~ (e20 = h5(e13)) & ~ (e20 = h5(e12)) & ~ (e20 = h5(e11)) & ~ (e20 = h5(e10))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))) & ((~ (e23 = h4(e13)) & ~ (e23 = h4(e12)) & ~ (e23 = h4(e11)) & ~ (e23 = h4(e10))) | (~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10))) | (~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10))) | (~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10))) | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) & ((~ (e23 = h3(e13)) & ~ (e23 = h3(e12)) & ~ (e23 = h3(e11)) & ~ (e23 = h3(e10))) | (~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10))) | (~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10))) | (~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10))) | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) & ((~ (e23 = h2(e13)) & ~ (e23 = h2(e12)) & ~ (e23 = h2(e11)) & ~ (e23 = h2(e10))) | (~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10))) | (~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10))) | (~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) & ((~ (e23 = h1(e13)) & ~ (e23 = h1(e12)) & ~ (e23 = h1(e11)) & ~ (e23 = h1(e10))) | (~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10))) | (~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10))) | (~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10))) | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(ennf_transformation, [], [f39])).
fof(f39, plain, ~ ((((e23 = h12(e13)) | (e23 = h12(e12)) | (e23 = h12(e11)) | (e23 = h12(e10))) & ((e22 = h12(e13)) | (e22 = h12(e12)) | (e22 = h12(e11)) | (e22 = h12(e10))) & ((e21 = h12(e13)) | (e21 = h12(e12)) | (e21 = h12(e11)) | (e21 = h12(e10))) & ((e20 = h12(e13)) | (e20 = h12(e12)) | (e20 = h12(e11)) | (e20 = h12(e10))) & (h12(op1(e13, e13)) = op2(h12(e13), h12(e13))) & (h12(op1(e13, e12)) = op2(h12(e13), h12(e12))) & (h12(op1(e13, e11)) = op2(h12(e13), h12(e11))) & (h12(op1(e13, e10)) = op2(h12(e13), h12(e10))) & (h12(op1(e12, e13)) = op2(h12(e12), h12(e13))) & (h12(op1(e12, e12)) = op2(h12(e12), h12(e12))) & (h12(op1(e12, e11)) = op2(h12(e12), h12(e11))) & (h12(op1(e12, e10)) = op2(h12(e12), h12(e10))) & (h12(op1(e11, e13)) = op2(h12(e11), h12(e13))) & (h12(op1(e11, e12)) = op2(h12(e11), h12(e12))) & (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) & (h12(op1(e11, e10)) = op2(h12(e11), h12(e10))) & (h12(op1(e10, e13)) = op2(h12(e10), h12(e13))) & (h12(op1(e10, e12)) = op2(h12(e10), h12(e12))) & (h12(op1(e10, e11)) = op2(h12(e10), h12(e11))) & (h12(op1(e10, e10)) = op2(h12(e10), h12(e10)))) | (((e23 = h11(e13)) | (e23 = h11(e12)) | (e23 = h11(e11)) | (e23 = h11(e10))) & ((e22 = h11(e13)) | (e22 = h11(e12)) | (e22 = h11(e11)) | (e22 = h11(e10))) & ((e21 = h11(e13)) | (e21 = h11(e12)) | (e21 = h11(e11)) | (e21 = h11(e10))) & ((e20 = h11(e13)) | (e20 = h11(e12)) | (e20 = h11(e11)) | (e20 = h11(e10))) & (h11(op1(e13, e13)) = op2(h11(e13), h11(e13))) & (h11(op1(e13, e12)) = op2(h11(e13), h11(e12))) & (h11(op1(e13, e11)) = op2(h11(e13), h11(e11))) & (h11(op1(e13, e10)) = op2(h11(e13), h11(e10))) & (h11(op1(e12, e13)) = op2(h11(e12), h11(e13))) & (h11(op1(e12, e12)) = op2(h11(e12), h11(e12))) & (h11(op1(e12, e11)) = op2(h11(e12), h11(e11))) & (h11(op1(e12, e10)) = op2(h11(e12), h11(e10))) & (h11(op1(e11, e13)) = op2(h11(e11), h11(e13))) & (h11(op1(e11, e12)) = op2(h11(e11), h11(e12))) & (h11(op1(e11, e11)) = op2(h11(e11), h11(e11))) & (h11(op1(e11, e10)) = op2(h11(e11), h11(e10))) & (h11(op1(e10, e13)) = op2(h11(e10), h11(e13))) & (h11(op1(e10, e12)) = op2(h11(e10), h11(e12))) & (h11(op1(e10, e11)) = op2(h11(e10), h11(e11))) & (h11(op1(e10, e10)) = op2(h11(e10), h11(e10)))) | (((e23 = h10(e13)) | (e23 = h10(e12)) | (e23 = h10(e11)) | (e23 = h10(e10))) & ((e22 = h10(e13)) | (e22 = h10(e12)) | (e22 = h10(e11)) | (e22 = h10(e10))) & ((e21 = h10(e13)) | (e21 = h10(e12)) | (e21 = h10(e11)) | (e21 = h10(e10))) & ((e20 = h10(e13)) | (e20 = h10(e12)) | (e20 = h10(e11)) | (e20 = h10(e10))) & (h10(op1(e13, e13)) = op2(h10(e13), h10(e13))) & (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) & (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) & (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) & (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) & (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) & (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) & (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) & (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) & (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) & (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) & (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) & (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) & (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) & (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) & (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))) | (((e23 = h9(e13)) | (e23 = h9(e12)) | (e23 = h9(e11)) | (e23 = h9(e10))) & ((e22 = h9(e13)) | (e22 = h9(e12)) | (e22 = h9(e11)) | (e22 = h9(e10))) & ((e21 = h9(e13)) | (e21 = h9(e12)) | (e21 = h9(e11)) | (e21 = h9(e10))) & ((e20 = h9(e13)) | (e20 = h9(e12)) | (e20 = h9(e11)) | (e20 = h9(e10))) & (h9(op1(e13, e13)) = op2(h9(e13), h9(e13))) & (h9(op1(e13, e12)) = op2(h9(e13), h9(e12))) & (h9(op1(e13, e11)) = op2(h9(e13), h9(e11))) & (h9(op1(e13, e10)) = op2(h9(e13), h9(e10))) & (h9(op1(e12, e13)) = op2(h9(e12), h9(e13))) & (h9(op1(e12, e12)) = op2(h9(e12), h9(e12))) & (h9(op1(e12, e11)) = op2(h9(e12), h9(e11))) & (h9(op1(e12, e10)) = op2(h9(e12), h9(e10))) & (h9(op1(e11, e13)) = op2(h9(e11), h9(e13))) & (h9(op1(e11, e12)) = op2(h9(e11), h9(e12))) & (h9(op1(e11, e11)) = op2(h9(e11), h9(e11))) & (h9(op1(e11, e10)) = op2(h9(e11), h9(e10))) & (h9(op1(e10, e13)) = op2(h9(e10), h9(e13))) & (h9(op1(e10, e12)) = op2(h9(e10), h9(e12))) & (h9(op1(e10, e11)) = op2(h9(e10), h9(e11))) & (h9(op1(e10, e10)) = op2(h9(e10), h9(e10)))) | (((e23 = h8(e13)) | (e23 = h8(e12)) | (e23 = h8(e11)) | (e23 = h8(e10))) & ((e22 = h8(e13)) | (e22 = h8(e12)) | (e22 = h8(e11)) | (e22 = h8(e10))) & ((e21 = h8(e13)) | (e21 = h8(e12)) | (e21 = h8(e11)) | (e21 = h8(e10))) & ((e20 = h8(e13)) | (e20 = h8(e12)) | (e20 = h8(e11)) | (e20 = h8(e10))) & (h8(op1(e13, e13)) = op2(h8(e13), h8(e13))) & (h8(op1(e13, e12)) = op2(h8(e13), h8(e12))) & (h8(op1(e13, e11)) = op2(h8(e13), h8(e11))) & (h8(op1(e13, e10)) = op2(h8(e13), h8(e10))) & (h8(op1(e12, e13)) = op2(h8(e12), h8(e13))) & (h8(op1(e12, e12)) = op2(h8(e12), h8(e12))) & (h8(op1(e12, e11)) = op2(h8(e12), h8(e11))) & (h8(op1(e12, e10)) = op2(h8(e12), h8(e10))) & (h8(op1(e11, e13)) = op2(h8(e11), h8(e13))) & (h8(op1(e11, e12)) = op2(h8(e11), h8(e12))) & (h8(op1(e11, e11)) = op2(h8(e11), h8(e11))) & (h8(op1(e11, e10)) = op2(h8(e11), h8(e10))) & (h8(op1(e10, e13)) = op2(h8(e10), h8(e13))) & (h8(op1(e10, e12)) = op2(h8(e10), h8(e12))) & (h8(op1(e10, e11)) = op2(h8(e10), h8(e11))) & (h8(op1(e10, e10)) = op2(h8(e10), h8(e10)))) | (((e23 = h7(e13)) | (e23 = h7(e12)) | (e23 = h7(e11)) | (e23 = h7(e10))) & ((e22 = h7(e13)) | (e22 = h7(e12)) | (e22 = h7(e11)) | (e22 = h7(e10))) & ((e21 = h7(e13)) | (e21 = h7(e12)) | (e21 = h7(e11)) | (e21 = h7(e10))) & ((e20 = h7(e13)) | (e20 = h7(e12)) | (e20 = h7(e11)) | (e20 = h7(e10))) & (h7(op1(e13, e13)) = op2(h7(e13), h7(e13))) & (h7(op1(e13, e12)) = op2(h7(e13), h7(e12))) & (h7(op1(e13, e11)) = op2(h7(e13), h7(e11))) & (h7(op1(e13, e10)) = op2(h7(e13), h7(e10))) & (h7(op1(e12, e13)) = op2(h7(e12), h7(e13))) & (h7(op1(e12, e12)) = op2(h7(e12), h7(e12))) & (h7(op1(e12, e11)) = op2(h7(e12), h7(e11))) & (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) & (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) & (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) & (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) & (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) & (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) & (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) & (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) & (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))) | (((e23 = h6(e13)) | (e23 = h6(e12)) | (e23 = h6(e11)) | (e23 = h6(e10))) & ((e22 = h6(e13)) | (e22 = h6(e12)) | (e22 = h6(e11)) | (e22 = h6(e10))) & ((e21 = h6(e13)) | (e21 = h6(e12)) | (e21 = h6(e11)) | (e21 = h6(e10))) & ((e20 = h6(e13)) | (e20 = h6(e12)) | (e20 = h6(e11)) | (e20 = h6(e10))) & (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) & (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) & (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) & (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) & (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) & (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) & (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) & (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) & (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) & (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) & (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) & (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) & (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) & (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) & (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) & (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))) | (((e23 = h5(e13)) | (e23 = h5(e12)) | (e23 = h5(e11)) | (e23 = h5(e10))) & ((e22 = h5(e13)) | (e22 = h5(e12)) | (e22 = h5(e11)) | (e22 = h5(e10))) & ((e21 = h5(e13)) | (e21 = h5(e12)) | (e21 = h5(e11)) | (e21 = h5(e10))) & ((e20 = h5(e13)) | (e20 = h5(e12)) | (e20 = h5(e11)) | (e20 = h5(e10))) & (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) & (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) & (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) & (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) & (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) & (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) & (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) & (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) & (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) & (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) & (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) & (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) & (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) & (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) & (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) & (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))) | (((e23 = h4(e13)) | (e23 = h4(e12)) | (e23 = h4(e11)) | (e23 = h4(e10))) & ((e22 = h4(e13)) | (e22 = h4(e12)) | (e22 = h4(e11)) | (e22 = h4(e10))) & ((e21 = h4(e13)) | (e21 = h4(e12)) | (e21 = h4(e11)) | (e21 = h4(e10))) & ((e20 = h4(e13)) | (e20 = h4(e12)) | (e20 = h4(e11)) | (e20 = h4(e10))) & (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) & (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) & (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) & (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) & (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) & (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) & (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) & (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) & (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) & (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) & (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) & (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) & (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) & (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) & (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) & (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) | (((e23 = h3(e13)) | (e23 = h3(e12)) | (e23 = h3(e11)) | (e23 = h3(e10))) & ((e22 = h3(e13)) | (e22 = h3(e12)) | (e22 = h3(e11)) | (e22 = h3(e10))) & ((e21 = h3(e13)) | (e21 = h3(e12)) | (e21 = h3(e11)) | (e21 = h3(e10))) & ((e20 = h3(e13)) | (e20 = h3(e12)) | (e20 = h3(e11)) | (e20 = h3(e10))) & (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) & (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) & (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) & (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) & (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) & (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) & (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) & (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) & (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) & (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) & (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) & (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) & (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) & (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) & (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) & (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) | (((e23 = h2(e13)) | (e23 = h2(e12)) | (e23 = h2(e11)) | (e23 = h2(e10))) & ((e22 = h2(e13)) | (e22 = h2(e12)) | (e22 = h2(e11)) | (e22 = h2(e10))) & ((e21 = h2(e13)) | (e21 = h2(e12)) | (e21 = h2(e11)) | (e21 = h2(e10))) & ((e20 = h2(e13)) | (e20 = h2(e12)) | (e20 = h2(e11)) | (e20 = h2(e10))) & (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) & (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) & (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) & (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) & (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) & (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) & (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) & (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) & (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) & (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) & (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) & (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) & (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) & (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) & (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) & (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) | (((e23 = h1(e13)) | (e23 = h1(e12)) | (e23 = h1(e11)) | (e23 = h1(e10))) & ((e22 = h1(e13)) | (e22 = h1(e12)) | (e22 = h1(e11)) | (e22 = h1(e10))) & ((e21 = h1(e13)) | (e21 = h1(e12)) | (e21 = h1(e11)) | (e21 = h1(e10))) & ((e20 = h1(e13)) | (e20 = h1(e12)) | (e20 = h1(e11)) | (e20 = h1(e10))) & (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) & (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) & (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) & (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) & (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) & (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) & (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) & (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) & (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) & (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) & (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) & (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) & (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) & (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) & (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) & (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(negated_conjecture, [], [f38])).
fof(f38, plain, ~ ((((e23 = h12(e13)) | (e23 = h12(e12)) | (e23 = h12(e11)) | (e23 = h12(e10))) & ((e22 = h12(e13)) | (e22 = h12(e12)) | (e22 = h12(e11)) | (e22 = h12(e10))) & ((e21 = h12(e13)) | (e21 = h12(e12)) | (e21 = h12(e11)) | (e21 = h12(e10))) & ((e20 = h12(e13)) | (e20 = h12(e12)) | (e20 = h12(e11)) | (e20 = h12(e10))) & (h12(op1(e13, e13)) = op2(h12(e13), h12(e13))) & (h12(op1(e13, e12)) = op2(h12(e13), h12(e12))) & (h12(op1(e13, e11)) = op2(h12(e13), h12(e11))) & (h12(op1(e13, e10)) = op2(h12(e13), h12(e10))) & (h12(op1(e12, e13)) = op2(h12(e12), h12(e13))) & (h12(op1(e12, e12)) = op2(h12(e12), h12(e12))) & (h12(op1(e12, e11)) = op2(h12(e12), h12(e11))) & (h12(op1(e12, e10)) = op2(h12(e12), h12(e10))) & (h12(op1(e11, e13)) = op2(h12(e11), h12(e13))) & (h12(op1(e11, e12)) = op2(h12(e11), h12(e12))) & (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) & (h12(op1(e11, e10)) = op2(h12(e11), h12(e10))) & (h12(op1(e10, e13)) = op2(h12(e10), h12(e13))) & (h12(op1(e10, e12)) = op2(h12(e10), h12(e12))) & (h12(op1(e10, e11)) = op2(h12(e10), h12(e11))) & (h12(op1(e10, e10)) = op2(h12(e10), h12(e10)))) | (((e23 = h11(e13)) | (e23 = h11(e12)) | (e23 = h11(e11)) | (e23 = h11(e10))) & ((e22 = h11(e13)) | (e22 = h11(e12)) | (e22 = h11(e11)) | (e22 = h11(e10))) & ((e21 = h11(e13)) | (e21 = h11(e12)) | (e21 = h11(e11)) | (e21 = h11(e10))) & ((e20 = h11(e13)) | (e20 = h11(e12)) | (e20 = h11(e11)) | (e20 = h11(e10))) & (h11(op1(e13, e13)) = op2(h11(e13), h11(e13))) & (h11(op1(e13, e12)) = op2(h11(e13), h11(e12))) & (h11(op1(e13, e11)) = op2(h11(e13), h11(e11))) & (h11(op1(e13, e10)) = op2(h11(e13), h11(e10))) & (h11(op1(e12, e13)) = op2(h11(e12), h11(e13))) & (h11(op1(e12, e12)) = op2(h11(e12), h11(e12))) & (h11(op1(e12, e11)) = op2(h11(e12), h11(e11))) & (h11(op1(e12, e10)) = op2(h11(e12), h11(e10))) & (h11(op1(e11, e13)) = op2(h11(e11), h11(e13))) & (h11(op1(e11, e12)) = op2(h11(e11), h11(e12))) & (h11(op1(e11, e11)) = op2(h11(e11), h11(e11))) & (h11(op1(e11, e10)) = op2(h11(e11), h11(e10))) & (h11(op1(e10, e13)) = op2(h11(e10), h11(e13))) & (h11(op1(e10, e12)) = op2(h11(e10), h11(e12))) & (h11(op1(e10, e11)) = op2(h11(e10), h11(e11))) & (h11(op1(e10, e10)) = op2(h11(e10), h11(e10)))) | (((e23 = h10(e13)) | (e23 = h10(e12)) | (e23 = h10(e11)) | (e23 = h10(e10))) & ((e22 = h10(e13)) | (e22 = h10(e12)) | (e22 = h10(e11)) | (e22 = h10(e10))) & ((e21 = h10(e13)) | (e21 = h10(e12)) | (e21 = h10(e11)) | (e21 = h10(e10))) & ((e20 = h10(e13)) | (e20 = h10(e12)) | (e20 = h10(e11)) | (e20 = h10(e10))) & (h10(op1(e13, e13)) = op2(h10(e13), h10(e13))) & (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) & (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) & (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) & (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) & (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) & (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) & (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) & (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) & (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) & (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) & (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) & (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) & (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) & (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) & (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))) | (((e23 = h9(e13)) | (e23 = h9(e12)) | (e23 = h9(e11)) | (e23 = h9(e10))) & ((e22 = h9(e13)) | (e22 = h9(e12)) | (e22 = h9(e11)) | (e22 = h9(e10))) & ((e21 = h9(e13)) | (e21 = h9(e12)) | (e21 = h9(e11)) | (e21 = h9(e10))) & ((e20 = h9(e13)) | (e20 = h9(e12)) | (e20 = h9(e11)) | (e20 = h9(e10))) & (h9(op1(e13, e13)) = op2(h9(e13), h9(e13))) & (h9(op1(e13, e12)) = op2(h9(e13), h9(e12))) & (h9(op1(e13, e11)) = op2(h9(e13), h9(e11))) & (h9(op1(e13, e10)) = op2(h9(e13), h9(e10))) & (h9(op1(e12, e13)) = op2(h9(e12), h9(e13))) & (h9(op1(e12, e12)) = op2(h9(e12), h9(e12))) & (h9(op1(e12, e11)) = op2(h9(e12), h9(e11))) & (h9(op1(e12, e10)) = op2(h9(e12), h9(e10))) & (h9(op1(e11, e13)) = op2(h9(e11), h9(e13))) & (h9(op1(e11, e12)) = op2(h9(e11), h9(e12))) & (h9(op1(e11, e11)) = op2(h9(e11), h9(e11))) & (h9(op1(e11, e10)) = op2(h9(e11), h9(e10))) & (h9(op1(e10, e13)) = op2(h9(e10), h9(e13))) & (h9(op1(e10, e12)) = op2(h9(e10), h9(e12))) & (h9(op1(e10, e11)) = op2(h9(e10), h9(e11))) & (h9(op1(e10, e10)) = op2(h9(e10), h9(e10)))) | (((e23 = h8(e13)) | (e23 = h8(e12)) | (e23 = h8(e11)) | (e23 = h8(e10))) & ((e22 = h8(e13)) | (e22 = h8(e12)) | (e22 = h8(e11)) | (e22 = h8(e10))) & ((e21 = h8(e13)) | (e21 = h8(e12)) | (e21 = h8(e11)) | (e21 = h8(e10))) & ((e20 = h8(e13)) | (e20 = h8(e12)) | (e20 = h8(e11)) | (e20 = h8(e10))) & (h8(op1(e13, e13)) = op2(h8(e13), h8(e13))) & (h8(op1(e13, e12)) = op2(h8(e13), h8(e12))) & (h8(op1(e13, e11)) = op2(h8(e13), h8(e11))) & (h8(op1(e13, e10)) = op2(h8(e13), h8(e10))) & (h8(op1(e12, e13)) = op2(h8(e12), h8(e13))) & (h8(op1(e12, e12)) = op2(h8(e12), h8(e12))) & (h8(op1(e12, e11)) = op2(h8(e12), h8(e11))) & (h8(op1(e12, e10)) = op2(h8(e12), h8(e10))) & (h8(op1(e11, e13)) = op2(h8(e11), h8(e13))) & (h8(op1(e11, e12)) = op2(h8(e11), h8(e12))) & (h8(op1(e11, e11)) = op2(h8(e11), h8(e11))) & (h8(op1(e11, e10)) = op2(h8(e11), h8(e10))) & (h8(op1(e10, e13)) = op2(h8(e10), h8(e13))) & (h8(op1(e10, e12)) = op2(h8(e10), h8(e12))) & (h8(op1(e10, e11)) = op2(h8(e10), h8(e11))) & (h8(op1(e10, e10)) = op2(h8(e10), h8(e10)))) | (((e23 = h7(e13)) | (e23 = h7(e12)) | (e23 = h7(e11)) | (e23 = h7(e10))) & ((e22 = h7(e13)) | (e22 = h7(e12)) | (e22 = h7(e11)) | (e22 = h7(e10))) & ((e21 = h7(e13)) | (e21 = h7(e12)) | (e21 = h7(e11)) | (e21 = h7(e10))) & ((e20 = h7(e13)) | (e20 = h7(e12)) | (e20 = h7(e11)) | (e20 = h7(e10))) & (h7(op1(e13, e13)) = op2(h7(e13), h7(e13))) & (h7(op1(e13, e12)) = op2(h7(e13), h7(e12))) & (h7(op1(e13, e11)) = op2(h7(e13), h7(e11))) & (h7(op1(e13, e10)) = op2(h7(e13), h7(e10))) & (h7(op1(e12, e13)) = op2(h7(e12), h7(e13))) & (h7(op1(e12, e12)) = op2(h7(e12), h7(e12))) & (h7(op1(e12, e11)) = op2(h7(e12), h7(e11))) & (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) & (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) & (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) & (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) & (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) & (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) & (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) & (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) & (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))) | (((e23 = h6(e13)) | (e23 = h6(e12)) | (e23 = h6(e11)) | (e23 = h6(e10))) & ((e22 = h6(e13)) | (e22 = h6(e12)) | (e22 = h6(e11)) | (e22 = h6(e10))) & ((e21 = h6(e13)) | (e21 = h6(e12)) | (e21 = h6(e11)) | (e21 = h6(e10))) & ((e20 = h6(e13)) | (e20 = h6(e12)) | (e20 = h6(e11)) | (e20 = h6(e10))) & (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) & (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) & (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) & (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) & (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) & (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) & (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) & (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) & (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) & (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) & (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) & (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) & (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) & (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) & (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) & (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))) | (((e23 = h5(e13)) | (e23 = h5(e12)) | (e23 = h5(e11)) | (e23 = h5(e10))) & ((e22 = h5(e13)) | (e22 = h5(e12)) | (e22 = h5(e11)) | (e22 = h5(e10))) & ((e21 = h5(e13)) | (e21 = h5(e12)) | (e21 = h5(e11)) | (e21 = h5(e10))) & ((e20 = h5(e13)) | (e20 = h5(e12)) | (e20 = h5(e11)) | (e20 = h5(e10))) & (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) & (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) & (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) & (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) & (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) & (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) & (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) & (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) & (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) & (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) & (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) & (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) & (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) & (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) & (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) & (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))) | (((e23 = h4(e13)) | (e23 = h4(e12)) | (e23 = h4(e11)) | (e23 = h4(e10))) & ((e22 = h4(e13)) | (e22 = h4(e12)) | (e22 = h4(e11)) | (e22 = h4(e10))) & ((e21 = h4(e13)) | (e21 = h4(e12)) | (e21 = h4(e11)) | (e21 = h4(e10))) & ((e20 = h4(e13)) | (e20 = h4(e12)) | (e20 = h4(e11)) | (e20 = h4(e10))) & (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) & (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) & (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) & (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) & (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) & (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) & (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) & (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) & (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) & (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) & (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) & (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) & (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) & (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) & (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) & (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) | (((e23 = h3(e13)) | (e23 = h3(e12)) | (e23 = h3(e11)) | (e23 = h3(e10))) & ((e22 = h3(e13)) | (e22 = h3(e12)) | (e22 = h3(e11)) | (e22 = h3(e10))) & ((e21 = h3(e13)) | (e21 = h3(e12)) | (e21 = h3(e11)) | (e21 = h3(e10))) & ((e20 = h3(e13)) | (e20 = h3(e12)) | (e20 = h3(e11)) | (e20 = h3(e10))) & (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) & (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) & (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) & (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) & (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) & (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) & (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) & (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) & (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) & (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) & (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) & (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) & (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) & (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) & (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) & (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) | (((e23 = h2(e13)) | (e23 = h2(e12)) | (e23 = h2(e11)) | (e23 = h2(e10))) & ((e22 = h2(e13)) | (e22 = h2(e12)) | (e22 = h2(e11)) | (e22 = h2(e10))) & ((e21 = h2(e13)) | (e21 = h2(e12)) | (e21 = h2(e11)) | (e21 = h2(e10))) & ((e20 = h2(e13)) | (e20 = h2(e12)) | (e20 = h2(e11)) | (e20 = h2(e10))) & (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) & (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) & (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) & (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) & (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) & (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) & (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) & (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) & (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) & (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) & (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) & (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) & (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) & (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) & (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) & (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) | (((e23 = h1(e13)) | (e23 = h1(e12)) | (e23 = h1(e11)) | (e23 = h1(e10))) & ((e22 = h1(e13)) | (e22 = h1(e12)) | (e22 = h1(e11)) | (e22 = h1(e10))) & ((e21 = h1(e13)) | (e21 = h1(e12)) | (e21 = h1(e11)) | (e21 = h1(e10))) & ((e20 = h1(e13)) | (e20 = h1(e12)) | (e20 = h1(e11)) | (e20 = h1(e10))) & (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) & (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) & (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) & (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) & (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) & (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) & (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) & (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) & (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) & (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) & (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) & (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) & (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) & (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) & (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) & (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG015+1.p', co1)).
fof(f1802, plain, (~ spl42_178 | ~ spl42_180), inference(avatar_split_clause, [], [f1797, f1799, f1789])).
fof(f1797, plain, (~ (e20 = h2(e10)) | ~ sP39), inference(forward_demodulation, [], [f602, f1767])).
fof(f602, plain, (~ (e20 = h12(e10)) | ~ sP39), inference(cnf_transformation, [], [f96])).
fof(f96, plain, ((~ (e20 = h12(e13)) & ~ (e20 = h12(e12)) & ~ (e20 = h12(e11)) & ~ (e20 = h12(e10))) | ~ sP39), inference(nnf_transformation, [], [f84])).
fof(f1779, plain, (~ spl42_175 | ~ spl42_176), inference(avatar_split_clause, [], [f599, f1776, f1772])).
fof(f599, plain, (~ (e21 = h12(e11)) | ~ sP40), inference(cnf_transformation, [], [f95])).
fof(f95, plain, ((~ (e21 = h12(e13)) & ~ (e21 = h12(e12)) & ~ (e21 = h12(e11)) & ~ (e21 = h12(e10))) | ~ sP40), inference(nnf_transformation, [], [f85])).
fof(f1583, plain, spl42_85, inference(avatar_split_clause, [], [f544, f1144])).
fof(f544, plain, (e20 = op2(e23, e23)), inference(cnf_transformation, [], [f25])).
fof(f25, plain, ((e21 = op2(e22, e23)) & (e20 = op2(e23, e23))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG015+1.p', ax25)).
fof(f1582, plain, spl42_102, inference(avatar_split_clause, [], [f545, f1216])).
fof(f545, plain, (e21 = op2(e22, e23)), inference(cnf_transformation, [], [f25])).
fof(f1581, plain, spl42_1, inference(avatar_split_clause, [], [f542, f787])).
fof(f542, plain, (e10 = op1(e13, e13)), inference(cnf_transformation, [], [f24])).
fof(f24, plain, ((e11 = op1(e12, e13)) & (e10 = op1(e13, e13))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG015+1.p', ax24)).
fof(f1580, plain, spl42_18, inference(avatar_split_clause, [], [f543, f859])).
fof(f543, plain, (e11 = op1(e12, e13)), inference(cnf_transformation, [], [f24])).
fof(f1576, plain, (spl42_174 | spl42_173 | spl42_172 | spl42_88), inference(avatar_split_clause, [], [f541, f1156, f1553, f1561, f1569])).
fof(f1569, plain, (spl42_174 <=> sP3), introduced(avatar_definition, [new_symbols(naming, [spl42_174])])).
fof(f1561, plain, (spl42_173 <=> sP4), introduced(avatar_definition, [new_symbols(naming, [spl42_173])])).
fof(f1553, plain, (spl42_172 <=> sP5), introduced(avatar_definition, [new_symbols(naming, [spl42_172])])).
fof(f541, plain, ((e23 = op2(e23, e23)) | sP5 | sP4 | sP3), inference(cnf_transformation, [], [f50])).
fof(f50, plain, (((e23 = op2(e23, e23)) & (e23 = op2(e22, e22)) & (e23 = op2(e21, e21)) & (op2(e20, e20) = e23)) | sP5 | sP4 | sP3), inference(definition_folding, [], [f23, e49, e48, e47])).
fof(f47, plain, (((e20 = op2(e23, e23)) & (e20 = op2(e22, e22)) & (e20 = op2(e21, e21)) & (e20 = op2(e20, e20))) | ~ sP3), inference(usedef, [], [e47])).
fof(e47, plain, (sP3 <=> ((e20 = op2(e23, e23)) & (e20 = op2(e22, e22)) & (e20 = op2(e21, e21)) & (e20 = op2(e20, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f48, plain, (((e21 = op2(e23, e23)) & (e21 = op2(e22, e22)) & (e21 = op2(e21, e21)) & (op2(e20, e20) = e21)) | ~ sP4), inference(usedef, [], [e48])).
fof(e48, plain, (sP4 <=> ((e21 = op2(e23, e23)) & (e21 = op2(e22, e22)) & (e21 = op2(e21, e21)) & (op2(e20, e20) = e21))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f49, plain, (((e22 = op2(e23, e23)) & (e22 = op2(e22, e22)) & (e22 = op2(e21, e21)) & (op2(e20, e20) = e22)) | ~ sP5), inference(usedef, [], [e49])).
fof(e49, plain, (sP5 <=> ((e22 = op2(e23, e23)) & (e22 = op2(e22, e22)) & (e22 = op2(e21, e21)) & (op2(e20, e20) = e22))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f23, plain, (((e23 = op2(e23, e23)) & (e23 = op2(e22, e22)) & (e23 = op2(e21, e21)) & (op2(e20, e20) = e23)) | ((e22 = op2(e23, e23)) & (e22 = op2(e22, e22)) & (e22 = op2(e21, e21)) & (op2(e20, e20) = e22)) | ((e21 = op2(e23, e23)) & (e21 = op2(e22, e22)) & (e21 = op2(e21, e21)) & (op2(e20, e20) = e21)) | ((e20 = op2(e23, e23)) & (e20 = op2(e22, e22)) & (e20 = op2(e21, e21)) & (e20 = op2(e20, e20)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG015+1.p', ax23)).
fof(f1575, plain, (~ spl42_174 | spl42_145), inference(avatar_split_clause, [], [f534, f1399, f1569])).
fof(f534, plain, ((e20 = op2(e20, e20)) | ~ sP3), inference(cnf_transformation, [], [f93])).
fof(f93, plain, (((e20 = op2(e23, e23)) & (e20 = op2(e22, e22)) & (e20 = op2(e21, e21)) & (e20 = op2(e20, e20))) | ~ sP3), inference(nnf_transformation, [], [f47])).
fof(f1574, plain, (~ spl42_174 | spl42_125), inference(avatar_split_clause, [], [f535, f1314, f1569])).
fof(f535, plain, ((e20 = op2(e21, e21)) | ~ sP3), inference(cnf_transformation, [], [f93])).
fof(f1573, plain, (~ spl42_174 | spl42_105), inference(avatar_split_clause, [], [f536, f1229, f1569])).
fof(f536, plain, ((e20 = op2(e22, e22)) | ~ sP3), inference(cnf_transformation, [], [f93])).
fof(f1565, plain, (~ spl42_173 | spl42_106), inference(avatar_split_clause, [], [f532, f1233, f1561])).
fof(f532, plain, ((e21 = op2(e22, e22)) | ~ sP4), inference(cnf_transformation, [], [f92])).
fof(f92, plain, (((e21 = op2(e23, e23)) & (e21 = op2(e22, e22)) & (e21 = op2(e21, e21)) & (op2(e20, e20) = e21)) | ~ sP4), inference(nnf_transformation, [], [f48])).
fof(f1557, plain, (~ spl42_172 | spl42_107), inference(avatar_split_clause, [], [f528, f1237, f1553])).
fof(f528, plain, ((e22 = op2(e22, e22)) | ~ sP5), inference(cnf_transformation, [], [f91])).
fof(f91, plain, (((e22 = op2(e23, e23)) & (e22 = op2(e22, e22)) & (e22 = op2(e21, e21)) & (op2(e20, e20) = e22)) | ~ sP5), inference(nnf_transformation, [], [f49])).
fof(f1548, plain, (spl42_171 | spl42_170 | spl42_169 | spl42_4), inference(avatar_split_clause, [], [f525, f799, f1525, f1533, f1541])).
fof(f1541, plain, (spl42_171 <=> sP0), introduced(avatar_definition, [new_symbols(naming, [spl42_171])])).
fof(f1533, plain, (spl42_170 <=> sP1), introduced(avatar_definition, [new_symbols(naming, [spl42_170])])).
fof(f1525, plain, (spl42_169 <=> sP2), introduced(avatar_definition, [new_symbols(naming, [spl42_169])])).
fof(f525, plain, ((e13 = op1(e13, e13)) | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f46])).
fof(f46, plain, (((e13 = op1(e13, e13)) & (e13 = op1(e12, e12)) & (e13 = op1(e11, e11)) & (op1(e10, e10) = e13)) | sP2 | sP1 | sP0), inference(definition_folding, [], [f22, e45, e44, e43])).
fof(f43, plain, (((e10 = op1(e13, e13)) & (e10 = op1(e12, e12)) & (e10 = op1(e11, e11)) & (e10 = op1(e10, e10))) | ~ sP0), inference(usedef, [], [e43])).
fof(e43, plain, (sP0 <=> ((e10 = op1(e13, e13)) & (e10 = op1(e12, e12)) & (e10 = op1(e11, e11)) & (e10 = op1(e10, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f44, plain, (((e11 = op1(e13, e13)) & (e11 = op1(e12, e12)) & (e11 = op1(e11, e11)) & (op1(e10, e10) = e11)) | ~ sP1), inference(usedef, [], [e44])).
fof(e44, plain, (sP1 <=> ((e11 = op1(e13, e13)) & (e11 = op1(e12, e12)) & (e11 = op1(e11, e11)) & (op1(e10, e10) = e11))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f45, plain, (((e12 = op1(e13, e13)) & (e12 = op1(e12, e12)) & (e12 = op1(e11, e11)) & (op1(e10, e10) = e12)) | ~ sP2), inference(usedef, [], [e45])).
fof(e45, plain, (sP2 <=> ((e12 = op1(e13, e13)) & (e12 = op1(e12, e12)) & (e12 = op1(e11, e11)) & (op1(e10, e10) = e12))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f22, plain, (((e13 = op1(e13, e13)) & (e13 = op1(e12, e12)) & (e13 = op1(e11, e11)) & (op1(e10, e10) = e13)) | ((e12 = op1(e13, e13)) & (e12 = op1(e12, e12)) & (e12 = op1(e11, e11)) & (op1(e10, e10) = e12)) | ((e11 = op1(e13, e13)) & (e11 = op1(e12, e12)) & (e11 = op1(e11, e11)) & (op1(e10, e10) = e11)) | ((e10 = op1(e13, e13)) & (e10 = op1(e12, e12)) & (e10 = op1(e11, e11)) & (e10 = op1(e10, e10)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG015+1.p', ax22)).
fof(f1547, plain, (~ spl42_171 | spl42_61), inference(avatar_split_clause, [], [f518, f1042, f1541])).
fof(f518, plain, ((e10 = op1(e10, e10)) | ~ sP0), inference(cnf_transformation, [], [f90])).
fof(f90, plain, (((e10 = op1(e13, e13)) & (e10 = op1(e12, e12)) & (e10 = op1(e11, e11)) & (e10 = op1(e10, e10))) | ~ sP0), inference(nnf_transformation, [], [f43])).
fof(f1546, plain, (~ spl42_171 | spl42_41), inference(avatar_split_clause, [], [f519, f957, f1541])).
fof(f519, plain, ((e10 = op1(e11, e11)) | ~ sP0), inference(cnf_transformation, [], [f90])).
fof(f1545, plain, (~ spl42_171 | spl42_21), inference(avatar_split_clause, [], [f520, f872, f1541])).
fof(f520, plain, ((e10 = op1(e12, e12)) | ~ sP0), inference(cnf_transformation, [], [f90])).
fof(f1536, plain, (~ spl42_170 | spl42_2), inference(avatar_split_clause, [], [f517, f791, f1533])).
fof(f517, plain, ((e11 = op1(e13, e13)) | ~ sP1), inference(cnf_transformation, [], [f89])).
fof(f89, plain, (((e11 = op1(e13, e13)) & (e11 = op1(e12, e12)) & (e11 = op1(e11, e11)) & (op1(e10, e10) = e11)) | ~ sP1), inference(nnf_transformation, [], [f44])).
fof(f1528, plain, (~ spl42_169 | spl42_3), inference(avatar_split_clause, [], [f513, f795, f1525])).
fof(f513, plain, ((e12 = op1(e13, e13)) | ~ sP2), inference(cnf_transformation, [], [f88])).
fof(f88, plain, (((e12 = op1(e13, e13)) & (e12 = op1(e12, e12)) & (e12 = op1(e11, e11)) & (op1(e10, e10) = e12)) | ~ sP2), inference(nnf_transformation, [], [f45])).
fof(f1482, plain, (spl42_161 | spl42_162 | spl42_163 | spl42_164), inference(avatar_split_clause, [], [f329, f1479, f1475, f1471, f1467])).
fof(f329, plain, ((e23 = inv2(e21)) | (e22 = inv2(e21)) | (e21 = inv2(e21)) | (e20 = inv2(e21))), inference(cnf_transformation, [], [f8])).
fof(f1465, plain, (spl42_157 | spl42_158 | spl42_159 | spl42_160), inference(avatar_split_clause, [], [f330, f1462, f1458, f1454, f1450])).
fof(f330, plain, ((e23 = inv2(e22)) | (e22 = inv2(e22)) | (e21 = inv2(e22)) | (e20 = inv2(e22))), inference(cnf_transformation, [], [f8])).
fof(f1431, plain, (spl42_149 | spl42_150 | spl42_151 | spl42_152), inference(avatar_split_clause, [], [f319, f1428, f1424, f1420, f1416])).
fof(f319, plain, ((e23 = unit2) | (e22 = unit2) | (e21 = unit2) | (e20 = unit2)), inference(cnf_transformation, [], [f7])).
fof(f1193, plain, (spl42_93 | spl42_94 | spl42_95 | spl42_96), inference(avatar_split_clause, [], [f244, f1190, f1186, f1182, f1178])).
fof(f244, plain, ((e23 = op2(e23, e21)) | (e22 = op2(e23, e21)) | (e21 = op2(e23, e21)) | (e20 = op2(e23, e21))), inference(cnf_transformation, [], [f5])).
fof(f5, plain, (((e23 = op2(e23, e23)) | (e22 = op2(e23, e23)) | (e21 = op2(e23, e23)) | (e20 = op2(e23, e23))) & ((e23 = op2(e23, e22)) | (e22 = op2(e23, e22)) | (e21 = op2(e23, e22)) | (e20 = op2(e23, e22))) & ((e23 = op2(e23, e21)) | (e22 = op2(e23, e21)) | (e21 = op2(e23, e21)) | (e20 = op2(e23, e21))) & ((e23 = op2(e23, e20)) | (e22 = op2(e23, e20)) | (e21 = op2(e23, e20)) | (e20 = op2(e23, e20))) & ((e23 = op2(e22, e23)) | (e22 = op2(e22, e23)) | (e21 = op2(e22, e23)) | (e20 = op2(e22, e23))) & ((e23 = op2(e22, e22)) | (e22 = op2(e22, e22)) | (e21 = op2(e22, e22)) | (e20 = op2(e22, e22))) & ((e23 = op2(e22, e21)) | (e22 = op2(e22, e21)) | (e21 = op2(e22, e21)) | (e20 = op2(e22, e21))) & ((e23 = op2(e22, e20)) | (e22 = op2(e22, e20)) | (e21 = op2(e22, e20)) | (e20 = op2(e22, e20))) & ((e23 = op2(e21, e23)) | (e22 = op2(e21, e23)) | (e21 = op2(e21, e23)) | (e20 = op2(e21, e23))) & ((e23 = op2(e21, e22)) | (e22 = op2(e21, e22)) | (e21 = op2(e21, e22)) | (e20 = op2(e21, e22))) & ((e23 = op2(e21, e21)) | (e22 = op2(e21, e21)) | (e21 = op2(e21, e21)) | (e20 = op2(e21, e21))) & ((e23 = op2(e21, e20)) | (e22 = op2(e21, e20)) | (e21 = op2(e21, e20)) | (e20 = op2(e21, e20))) & ((e23 = op2(e20, e23)) | (e22 = op2(e20, e23)) | (e21 = op2(e20, e23)) | (e20 = op2(e20, e23))) & ((e23 = op2(e20, e22)) | (e22 = op2(e20, e22)) | (e21 = op2(e20, e22)) | (e20 = op2(e20, e22))) & ((e23 = op2(e20, e21)) | (e22 = op2(e20, e21)) | (e21 = op2(e20, e21)) | (e20 = op2(e20, e21))) & ((op2(e20, e20) = e23) | (op2(e20, e20) = e22) | (op2(e20, e20) = e21) | (e20 = op2(e20, e20)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG015+1.p', ax5)).
fof(f1176, plain, (spl42_89 | spl42_90 | spl42_91 | spl42_92), inference(avatar_split_clause, [], [f245, f1173, f1169, f1165, f1161])).
fof(f245, plain, ((e23 = op2(e23, e22)) | (e22 = op2(e23, e22)) | (e21 = op2(e23, e22)) | (e20 = op2(e23, e22))), inference(cnf_transformation, [], [f5])).
fof(f1074, plain, (spl42_65 | spl42_66 | spl42_67 | spl42_68), inference(avatar_split_clause, [], [f218, f1071, f1067, f1063, f1059])).
fof(f218, plain, ((e13 = unit1) | (e12 = unit1) | (e11 = unit1) | (e10 = unit1)), inference(cnf_transformation, [], [f3])).