fof(f2863, plain, $false, inference(avatar_sat_refutation, [], [f1669, f2026, f2304, f2309, f2314, f2319, f2322, f2327, f2332, f2335, f2340, f2345, f2350, f2355, f2360, f2365, f2370, f2375, f2378, f2383, f2388, f2393, f2396, f2401, f2862])).
fof(f2862, plain, spl191_1, inference(avatar_split_clause, [], [f2861, f1647])).
fof(f1647, plain, (spl191_1 <=> sP115), introduced(avatar_definition, [new_symbols(naming, [spl191_1])])).
fof(f2861, plain, sP115, inference(subsumption_resolution, [], [f2860, f1645])).
fof(f1645, plain, ~ sP116, inference(subsumption_resolution, [], [f1270, f1087])).
fof(f1087, plain, (unit = op(unit, unit)), inference(definition_unfolding, [], [f400, f425, f425, f425])).
fof(f425, plain, (e0 = unit), inference(cnf_transformation, [], [f3])).
fof(f3, plain, (e0 = unit), file('/home/ubuntu/library/tptp/Problems/ALG/ALG093+1.p', ax3)).
fof(f400, plain, (e0 = op(e0, e0)), inference(cnf_transformation, [], [f2])).
fof(f2, plain, ((e2 = op(e4, e4)) & (e1 = op(e4, e3)) & (e0 = op(e4, e2)) & (e3 = op(e4, e1)) & (e4 = op(e4, e0)) & (e0 = op(e3, e4)) & (e4 = op(e3, e3)) & (e1 = op(e3, e2)) & (e2 = op(e3, e1)) & (e3 = op(e3, e0)) & (e1 = op(e2, e4)) & (e0 = op(e2, e3)) & (e3 = op(e2, e2)) & (e4 = op(e2, e1)) & (e2 = op(e2, e0)) & (e3 = op(e1, e4)) & (e2 = op(e1, e3)) & (e4 = op(e1, e2)) & (e0 = op(e1, e1)) & (e1 = op(e1, e0)) & (e4 = op(e0, e4)) & (e3 = op(e0, e3)) & (e2 = op(e0, e2)) & (e1 = op(e0, e1)) & (e0 = op(e0, e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG093+1.p', ax2)).
fof(f1270, plain, (~ (unit = op(unit, unit)) | ~ sP116), inference(definition_unfolding, [], [f796, f425, f425, f425])).
fof(f796, plain, (~ (e0 = op(e0, e0)) | ~ sP116), inference(cnf_transformation, [], [f273])).
fof(f273, plain, ((~ (e4 = op(e0, e0)) & ~ (e3 = op(e0, e0)) & ~ (e2 = op(e0, e0)) & ~ (e1 = op(e0, e0)) & ~ (e0 = op(e0, e0))) | ~ sP116), inference(nnf_transformation, [], [f123])).
fof(f123, plain, ((~ (e4 = op(e0, e0)) & ~ (e3 = op(e0, e0)) & ~ (e2 = op(e0, e0)) & ~ (e1 = op(e0, e0)) & ~ (e0 = op(e0, e0))) | ~ sP116), inference(usedef, [], [e123])).
fof(e123, plain, (sP116 <=> (~ (e4 = op(e0, e0)) & ~ (e3 = op(e0, e0)) & ~ (e2 = op(e0, e0)) & ~ (e1 = op(e0, e0)) & ~ (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP116])])).
fof(f2860, plain, (sP116 | sP115), inference(subsumption_resolution, [], [f2859, f1640])).
fof(f1640, plain, ~ sP117, inference(subsumption_resolution, [], [f1264, f1086])).
fof(f1086, plain, (e1 = op(unit, e1)), inference(definition_unfolding, [], [f401, f425])).
fof(f401, plain, (e1 = op(e0, e1)), inference(cnf_transformation, [], [f2])).
fof(f1264, plain, (~ (e1 = op(unit, e1)) | ~ sP117), inference(definition_unfolding, [], [f792, f425])).
fof(f792, plain, (~ (e1 = op(e0, e1)) | ~ sP117), inference(cnf_transformation, [], [f272])).
fof(f272, plain, ((~ (e4 = op(e0, e1)) & ~ (e3 = op(e0, e1)) & ~ (e2 = op(e0, e1)) & ~ (e1 = op(e0, e1)) & ~ (e0 = op(e0, e1))) | ~ sP117), inference(nnf_transformation, [], [f124])).
fof(f124, plain, ((~ (e4 = op(e0, e1)) & ~ (e3 = op(e0, e1)) & ~ (e2 = op(e0, e1)) & ~ (e1 = op(e0, e1)) & ~ (e0 = op(e0, e1))) | ~ sP117), inference(usedef, [], [e124])).
fof(e124, plain, (sP117 <=> (~ (e4 = op(e0, e1)) & ~ (e3 = op(e0, e1)) & ~ (e2 = op(e0, e1)) & ~ (e1 = op(e0, e1)) & ~ (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP117])])).
fof(f2859, plain, (sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2858, f1636])).
fof(f1636, plain, ~ sP118, inference(subsumption_resolution, [], [f1258, f1085])).
fof(f1085, plain, (e2 = op(unit, e2)), inference(definition_unfolding, [], [f402, f425])).
fof(f402, plain, (e2 = op(e0, e2)), inference(cnf_transformation, [], [f2])).
fof(f1258, plain, (~ (e2 = op(unit, e2)) | ~ sP118), inference(definition_unfolding, [], [f788, f425])).
fof(f788, plain, (~ (e2 = op(e0, e2)) | ~ sP118), inference(cnf_transformation, [], [f271])).
fof(f271, plain, ((~ (e4 = op(e0, e2)) & ~ (e3 = op(e0, e2)) & ~ (e2 = op(e0, e2)) & ~ (e1 = op(e0, e2)) & ~ (e0 = op(e0, e2))) | ~ sP118), inference(nnf_transformation, [], [f125])).
fof(f125, plain, ((~ (e4 = op(e0, e2)) & ~ (e3 = op(e0, e2)) & ~ (e2 = op(e0, e2)) & ~ (e1 = op(e0, e2)) & ~ (e0 = op(e0, e2))) | ~ sP118), inference(usedef, [], [e125])).
fof(e125, plain, (sP118 <=> (~ (e4 = op(e0, e2)) & ~ (e3 = op(e0, e2)) & ~ (e2 = op(e0, e2)) & ~ (e1 = op(e0, e2)) & ~ (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP118])])).
fof(f2858, plain, (sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2857, f1633])).
fof(f1633, plain, ~ sP119, inference(subsumption_resolution, [], [f1252, f1084])).
fof(f1084, plain, (e3 = op(unit, e3)), inference(definition_unfolding, [], [f403, f425])).
fof(f403, plain, (e3 = op(e0, e3)), inference(cnf_transformation, [], [f2])).
fof(f1252, plain, (~ (e3 = op(unit, e3)) | ~ sP119), inference(definition_unfolding, [], [f784, f425])).
fof(f784, plain, (~ (e3 = op(e0, e3)) | ~ sP119), inference(cnf_transformation, [], [f270])).
fof(f270, plain, ((~ (e4 = op(e0, e3)) & ~ (e3 = op(e0, e3)) & ~ (e2 = op(e0, e3)) & ~ (e1 = op(e0, e3)) & ~ (e0 = op(e0, e3))) | ~ sP119), inference(nnf_transformation, [], [f126])).
fof(f126, plain, ((~ (e4 = op(e0, e3)) & ~ (e3 = op(e0, e3)) & ~ (e2 = op(e0, e3)) & ~ (e1 = op(e0, e3)) & ~ (e0 = op(e0, e3))) | ~ sP119), inference(usedef, [], [e126])).
fof(e126, plain, (sP119 <=> (~ (e4 = op(e0, e3)) & ~ (e3 = op(e0, e3)) & ~ (e2 = op(e0, e3)) & ~ (e1 = op(e0, e3)) & ~ (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP119])])).
fof(f2857, plain, (sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2856, f1631])).
fof(f1631, plain, ~ sP120, inference(subsumption_resolution, [], [f1246, f1083])).
fof(f1083, plain, (e4 = op(unit, e4)), inference(definition_unfolding, [], [f404, f425])).
fof(f404, plain, (e4 = op(e0, e4)), inference(cnf_transformation, [], [f2])).
fof(f1246, plain, (~ (e4 = op(unit, e4)) | ~ sP120), inference(definition_unfolding, [], [f780, f425])).
fof(f780, plain, (~ (e4 = op(e0, e4)) | ~ sP120), inference(cnf_transformation, [], [f269])).
fof(f269, plain, ((~ (e4 = op(e0, e4)) & ~ (e3 = op(e0, e4)) & ~ (e2 = op(e0, e4)) & ~ (e1 = op(e0, e4)) & ~ (e0 = op(e0, e4))) | ~ sP120), inference(nnf_transformation, [], [f127])).
fof(f127, plain, ((~ (e4 = op(e0, e4)) & ~ (e3 = op(e0, e4)) & ~ (e2 = op(e0, e4)) & ~ (e1 = op(e0, e4)) & ~ (e0 = op(e0, e4))) | ~ sP120), inference(usedef, [], [e127])).
fof(e127, plain, (sP120 <=> (~ (e4 = op(e0, e4)) & ~ (e3 = op(e0, e4)) & ~ (e2 = op(e0, e4)) & ~ (e1 = op(e0, e4)) & ~ (e0 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP120])])).
fof(f2856, plain, (sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2855, f1630])).
fof(f1630, plain, ~ sP121, inference(subsumption_resolution, [], [f1244, f1082])).
fof(f1082, plain, (e1 = op(e1, unit)), inference(definition_unfolding, [], [f405, f425])).
fof(f405, plain, (e1 = op(e1, e0)), inference(cnf_transformation, [], [f2])).
fof(f1244, plain, (~ (e1 = op(e1, unit)) | ~ sP121), inference(definition_unfolding, [], [f772, f425])).
fof(f772, plain, (~ (e1 = op(e1, e0)) | ~ sP121), inference(cnf_transformation, [], [f268])).
fof(f268, plain, ((~ (e4 = op(e1, e0)) & ~ (e3 = op(e1, e0)) & ~ (e2 = op(e1, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e1, e0))) | ~ sP121), inference(nnf_transformation, [], [f128])).
fof(f128, plain, ((~ (e4 = op(e1, e0)) & ~ (e3 = op(e1, e0)) & ~ (e2 = op(e1, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e1, e0))) | ~ sP121), inference(usedef, [], [e128])).
fof(e128, plain, (sP121 <=> (~ (e4 = op(e1, e0)) & ~ (e3 = op(e1, e0)) & ~ (e2 = op(e1, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP121])])).
fof(f2855, plain, (sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2854, f1626])).
fof(f1626, plain, ~ sP122, inference(subsumption_resolution, [], [f1240, f1081])).
fof(f1081, plain, (op(e1, e1) = unit), inference(definition_unfolding, [], [f406, f425])).
fof(f406, plain, (e0 = op(e1, e1)), inference(cnf_transformation, [], [f2])).
fof(f1240, plain, (~ (op(e1, e1) = unit) | ~ sP122), inference(definition_unfolding, [], [f766, f425])).
fof(f766, plain, (~ (e0 = op(e1, e1)) | ~ sP122), inference(cnf_transformation, [], [f267])).
fof(f267, plain, ((~ (e4 = op(e1, e1)) & ~ (e3 = op(e1, e1)) & ~ (e2 = op(e1, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e1, e1))) | ~ sP122), inference(nnf_transformation, [], [f129])).
fof(f129, plain, ((~ (e4 = op(e1, e1)) & ~ (e3 = op(e1, e1)) & ~ (e2 = op(e1, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e1, e1))) | ~ sP122), inference(usedef, [], [e129])).
fof(e129, plain, (sP122 <=> (~ (e4 = op(e1, e1)) & ~ (e3 = op(e1, e1)) & ~ (e2 = op(e1, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP122])])).
fof(f2854, plain, (sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2853, f1621])).
fof(f1621, plain, ~ sP123, inference(subsumption_resolution, [], [f765, f407])).
fof(f407, plain, (e4 = op(e1, e2)), inference(cnf_transformation, [], [f2])).
fof(f765, plain, (~ (e4 = op(e1, e2)) | ~ sP123), inference(cnf_transformation, [], [f266])).
fof(f266, plain, ((~ (e4 = op(e1, e2)) & ~ (e3 = op(e1, e2)) & ~ (e2 = op(e1, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e1, e2))) | ~ sP123), inference(nnf_transformation, [], [f130])).
fof(f130, plain, ((~ (e4 = op(e1, e2)) & ~ (e3 = op(e1, e2)) & ~ (e2 = op(e1, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e1, e2))) | ~ sP123), inference(usedef, [], [e130])).
fof(e130, plain, (sP123 <=> (~ (e4 = op(e1, e2)) & ~ (e3 = op(e1, e2)) & ~ (e2 = op(e1, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP123])])).
fof(f2853, plain, (sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2852, f1620])).
fof(f1620, plain, ~ sP124, inference(subsumption_resolution, [], [f758, f408])).
fof(f408, plain, (e2 = op(e1, e3)), inference(cnf_transformation, [], [f2])).
fof(f758, plain, (~ (e2 = op(e1, e3)) | ~ sP124), inference(cnf_transformation, [], [f265])).
fof(f265, plain, ((~ (e4 = op(e1, e3)) & ~ (e3 = op(e1, e3)) & ~ (e2 = op(e1, e3)) & ~ (e1 = op(e1, e3)) & ~ (e0 = op(e1, e3))) | ~ sP124), inference(nnf_transformation, [], [f131])).
fof(f131, plain, ((~ (e4 = op(e1, e3)) & ~ (e3 = op(e1, e3)) & ~ (e2 = op(e1, e3)) & ~ (e1 = op(e1, e3)) & ~ (e0 = op(e1, e3))) | ~ sP124), inference(usedef, [], [e131])).
fof(e131, plain, (sP124 <=> (~ (e4 = op(e1, e3)) & ~ (e3 = op(e1, e3)) & ~ (e2 = op(e1, e3)) & ~ (e1 = op(e1, e3)) & ~ (e0 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP124])])).
fof(f2852, plain, (sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2851, f1617])).
fof(f1617, plain, ~ sP125, inference(subsumption_resolution, [], [f754, f409])).
fof(f409, plain, (e3 = op(e1, e4)), inference(cnf_transformation, [], [f2])).
fof(f754, plain, (~ (e3 = op(e1, e4)) | ~ sP125), inference(cnf_transformation, [], [f264])).
fof(f264, plain, ((~ (e4 = op(e1, e4)) & ~ (e3 = op(e1, e4)) & ~ (e2 = op(e1, e4)) & ~ (e1 = op(e1, e4)) & ~ (e0 = op(e1, e4))) | ~ sP125), inference(nnf_transformation, [], [f132])).
fof(f132, plain, ((~ (e4 = op(e1, e4)) & ~ (e3 = op(e1, e4)) & ~ (e2 = op(e1, e4)) & ~ (e1 = op(e1, e4)) & ~ (e0 = op(e1, e4))) | ~ sP125), inference(usedef, [], [e132])).
fof(e132, plain, (sP125 <=> (~ (e4 = op(e1, e4)) & ~ (e3 = op(e1, e4)) & ~ (e2 = op(e1, e4)) & ~ (e1 = op(e1, e4)) & ~ (e0 = op(e1, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP125])])).
fof(f2851, plain, (sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2850, f1615])).
fof(f1615, plain, ~ sP126, inference(subsumption_resolution, [], [f1234, f1080])).
fof(f1080, plain, (e2 = op(e2, unit)), inference(definition_unfolding, [], [f410, f425])).
fof(f410, plain, (e2 = op(e2, e0)), inference(cnf_transformation, [], [f2])).
fof(f1234, plain, (~ (e2 = op(e2, unit)) | ~ sP126), inference(definition_unfolding, [], [f748, f425])).
fof(f748, plain, (~ (e2 = op(e2, e0)) | ~ sP126), inference(cnf_transformation, [], [f263])).
fof(f263, plain, ((~ (e4 = op(e2, e0)) & ~ (e3 = op(e2, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e2, e0)) & ~ (e0 = op(e2, e0))) | ~ sP126), inference(nnf_transformation, [], [f133])).
fof(f133, plain, ((~ (e4 = op(e2, e0)) & ~ (e3 = op(e2, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e2, e0)) & ~ (e0 = op(e2, e0))) | ~ sP126), inference(usedef, [], [e133])).
fof(e133, plain, (sP126 <=> (~ (e4 = op(e2, e0)) & ~ (e3 = op(e2, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e2, e0)) & ~ (e0 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP126])])).
fof(f2850, plain, (sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2849, f1612])).
fof(f1612, plain, ~ sP127, inference(subsumption_resolution, [], [f745, f411])).
fof(f411, plain, (e4 = op(e2, e1)), inference(cnf_transformation, [], [f2])).
fof(f745, plain, (~ (e4 = op(e2, e1)) | ~ sP127), inference(cnf_transformation, [], [f262])).
fof(f262, plain, ((~ (e4 = op(e2, e1)) & ~ (e3 = op(e2, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e2, e1)) & ~ (e0 = op(e2, e1))) | ~ sP127), inference(nnf_transformation, [], [f134])).
fof(f134, plain, ((~ (e4 = op(e2, e1)) & ~ (e3 = op(e2, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e2, e1)) & ~ (e0 = op(e2, e1))) | ~ sP127), inference(usedef, [], [e134])).
fof(e134, plain, (sP127 <=> (~ (e4 = op(e2, e1)) & ~ (e3 = op(e2, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e2, e1)) & ~ (e0 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP127])])).
fof(f2849, plain, (sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2848, f1611])).
fof(f1611, plain, ~ sP128, inference(subsumption_resolution, [], [f739, f412])).
fof(f412, plain, (e3 = op(e2, e2)), inference(cnf_transformation, [], [f2])).
fof(f739, plain, (~ (e3 = op(e2, e2)) | ~ sP128), inference(cnf_transformation, [], [f261])).
fof(f261, plain, ((~ (e4 = op(e2, e2)) & ~ (e3 = op(e2, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e2, e2)) & ~ (e0 = op(e2, e2))) | ~ sP128), inference(nnf_transformation, [], [f135])).
fof(f135, plain, ((~ (e4 = op(e2, e2)) & ~ (e3 = op(e2, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e2, e2)) & ~ (e0 = op(e2, e2))) | ~ sP128), inference(usedef, [], [e135])).
fof(e135, plain, (sP128 <=> (~ (e4 = op(e2, e2)) & ~ (e3 = op(e2, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e2, e2)) & ~ (e0 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP128])])).
fof(f2848, plain, (sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2847, f1609])).
fof(f1609, plain, ~ sP129, inference(subsumption_resolution, [], [f1229, f1079])).
fof(f1079, plain, (op(e2, e3) = unit), inference(definition_unfolding, [], [f413, f425])).
fof(f413, plain, (e0 = op(e2, e3)), inference(cnf_transformation, [], [f2])).
fof(f1229, plain, (~ (op(e2, e3) = unit) | ~ sP129), inference(definition_unfolding, [], [f731, f425])).
fof(f731, plain, (~ (e0 = op(e2, e3)) | ~ sP129), inference(cnf_transformation, [], [f260])).
fof(f260, plain, ((~ (e4 = op(e2, e3)) & ~ (e3 = op(e2, e3)) & ~ (e2 = op(e2, e3)) & ~ (e1 = op(e2, e3)) & ~ (e0 = op(e2, e3))) | ~ sP129), inference(nnf_transformation, [], [f136])).
fof(f136, plain, ((~ (e4 = op(e2, e3)) & ~ (e3 = op(e2, e3)) & ~ (e2 = op(e2, e3)) & ~ (e1 = op(e2, e3)) & ~ (e0 = op(e2, e3))) | ~ sP129), inference(usedef, [], [e136])).
fof(e136, plain, (sP129 <=> (~ (e4 = op(e2, e3)) & ~ (e3 = op(e2, e3)) & ~ (e2 = op(e2, e3)) & ~ (e1 = op(e2, e3)) & ~ (e0 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP129])])).
fof(f2847, plain, (sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2846, f1604])).
fof(f1604, plain, ~ sP130, inference(subsumption_resolution, [], [f727, f414])).
fof(f414, plain, (e1 = op(e2, e4)), inference(cnf_transformation, [], [f2])).
fof(f727, plain, (~ (e1 = op(e2, e4)) | ~ sP130), inference(cnf_transformation, [], [f259])).
fof(f259, plain, ((~ (e4 = op(e2, e4)) & ~ (e3 = op(e2, e4)) & ~ (e2 = op(e2, e4)) & ~ (e1 = op(e2, e4)) & ~ (e0 = op(e2, e4))) | ~ sP130), inference(nnf_transformation, [], [f137])).
fof(f137, plain, ((~ (e4 = op(e2, e4)) & ~ (e3 = op(e2, e4)) & ~ (e2 = op(e2, e4)) & ~ (e1 = op(e2, e4)) & ~ (e0 = op(e2, e4))) | ~ sP130), inference(usedef, [], [e137])).
fof(e137, plain, (sP130 <=> (~ (e4 = op(e2, e4)) & ~ (e3 = op(e2, e4)) & ~ (e2 = op(e2, e4)) & ~ (e1 = op(e2, e4)) & ~ (e0 = op(e2, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP130])])).
fof(f2846, plain, (sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2845, f1600])).
fof(f1600, plain, ~ sP131, inference(subsumption_resolution, [], [f1224, f1078])).
fof(f1078, plain, (e3 = op(e3, unit)), inference(definition_unfolding, [], [f415, f425])).
fof(f415, plain, (e3 = op(e3, e0)), inference(cnf_transformation, [], [f2])).
fof(f1224, plain, (~ (e3 = op(e3, unit)) | ~ sP131), inference(definition_unfolding, [], [f724, f425])).
fof(f724, plain, (~ (e3 = op(e3, e0)) | ~ sP131), inference(cnf_transformation, [], [f258])).
fof(f258, plain, ((~ (e4 = op(e3, e0)) & ~ (e3 = op(e3, e0)) & ~ (e2 = op(e3, e0)) & ~ (e1 = op(e3, e0)) & ~ (e0 = op(e3, e0))) | ~ sP131), inference(nnf_transformation, [], [f138])).
fof(f138, plain, ((~ (e4 = op(e3, e0)) & ~ (e3 = op(e3, e0)) & ~ (e2 = op(e3, e0)) & ~ (e1 = op(e3, e0)) & ~ (e0 = op(e3, e0))) | ~ sP131), inference(usedef, [], [e138])).
fof(e138, plain, (sP131 <=> (~ (e4 = op(e3, e0)) & ~ (e3 = op(e3, e0)) & ~ (e2 = op(e3, e0)) & ~ (e1 = op(e3, e0)) & ~ (e0 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP131])])).
fof(f2845, plain, (sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2844, f1598])).
fof(f1598, plain, ~ sP132, inference(subsumption_resolution, [], [f718, f416])).
fof(f416, plain, (e2 = op(e3, e1)), inference(cnf_transformation, [], [f2])).
fof(f718, plain, (~ (e2 = op(e3, e1)) | ~ sP132), inference(cnf_transformation, [], [f257])).
fof(f257, plain, ((~ (e4 = op(e3, e1)) & ~ (e3 = op(e3, e1)) & ~ (e2 = op(e3, e1)) & ~ (e1 = op(e3, e1)) & ~ (e0 = op(e3, e1))) | ~ sP132), inference(nnf_transformation, [], [f139])).
fof(f139, plain, ((~ (e4 = op(e3, e1)) & ~ (e3 = op(e3, e1)) & ~ (e2 = op(e3, e1)) & ~ (e1 = op(e3, e1)) & ~ (e0 = op(e3, e1))) | ~ sP132), inference(usedef, [], [e139])).
fof(e139, plain, (sP132 <=> (~ (e4 = op(e3, e1)) & ~ (e3 = op(e3, e1)) & ~ (e2 = op(e3, e1)) & ~ (e1 = op(e3, e1)) & ~ (e0 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP132])])).
fof(f2844, plain, (sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2843, f1595])).
fof(f1595, plain, ~ sP133, inference(subsumption_resolution, [], [f712, f417])).
fof(f417, plain, (e1 = op(e3, e2)), inference(cnf_transformation, [], [f2])).
fof(f712, plain, (~ (e1 = op(e3, e2)) | ~ sP133), inference(cnf_transformation, [], [f256])).
fof(f256, plain, ((~ (e4 = op(e3, e2)) & ~ (e3 = op(e3, e2)) & ~ (e2 = op(e3, e2)) & ~ (e1 = op(e3, e2)) & ~ (e0 = op(e3, e2))) | ~ sP133), inference(nnf_transformation, [], [f140])).
fof(f140, plain, ((~ (e4 = op(e3, e2)) & ~ (e3 = op(e3, e2)) & ~ (e2 = op(e3, e2)) & ~ (e1 = op(e3, e2)) & ~ (e0 = op(e3, e2))) | ~ sP133), inference(usedef, [], [e140])).
fof(e140, plain, (sP133 <=> (~ (e4 = op(e3, e2)) & ~ (e3 = op(e3, e2)) & ~ (e2 = op(e3, e2)) & ~ (e1 = op(e3, e2)) & ~ (e0 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP133])])).
fof(f2843, plain, (sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2842, f1591])).
fof(f1591, plain, ~ sP134, inference(subsumption_resolution, [], [f710, f418])).
fof(f418, plain, (e4 = op(e3, e3)), inference(cnf_transformation, [], [f2])).
fof(f710, plain, (~ (e4 = op(e3, e3)) | ~ sP134), inference(cnf_transformation, [], [f255])).
fof(f255, plain, ((~ (e4 = op(e3, e3)) & ~ (e3 = op(e3, e3)) & ~ (e2 = op(e3, e3)) & ~ (e1 = op(e3, e3)) & ~ (e0 = op(e3, e3))) | ~ sP134), inference(nnf_transformation, [], [f141])).
fof(f141, plain, ((~ (e4 = op(e3, e3)) & ~ (e3 = op(e3, e3)) & ~ (e2 = op(e3, e3)) & ~ (e1 = op(e3, e3)) & ~ (e0 = op(e3, e3))) | ~ sP134), inference(usedef, [], [e141])).
fof(e141, plain, (sP134 <=> (~ (e4 = op(e3, e3)) & ~ (e3 = op(e3, e3)) & ~ (e2 = op(e3, e3)) & ~ (e1 = op(e3, e3)) & ~ (e0 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP134])])).
fof(f2842, plain, (sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2841, f1590])).
fof(f1590, plain, ~ sP135, inference(subsumption_resolution, [], [f1219, f1077])).
fof(f1077, plain, (op(e3, e4) = unit), inference(definition_unfolding, [], [f419, f425])).
fof(f419, plain, (e0 = op(e3, e4)), inference(cnf_transformation, [], [f2])).
fof(f1219, plain, (~ (op(e3, e4) = unit) | ~ sP135), inference(definition_unfolding, [], [f701, f425])).
fof(f701, plain, (~ (e0 = op(e3, e4)) | ~ sP135), inference(cnf_transformation, [], [f254])).
fof(f254, plain, ((~ (e4 = op(e3, e4)) & ~ (e3 = op(e3, e4)) & ~ (e2 = op(e3, e4)) & ~ (e1 = op(e3, e4)) & ~ (e0 = op(e3, e4))) | ~ sP135), inference(nnf_transformation, [], [f142])).
fof(f142, plain, ((~ (e4 = op(e3, e4)) & ~ (e3 = op(e3, e4)) & ~ (e2 = op(e3, e4)) & ~ (e1 = op(e3, e4)) & ~ (e0 = op(e3, e4))) | ~ sP135), inference(usedef, [], [e142])).
fof(e142, plain, (sP135 <=> (~ (e4 = op(e3, e4)) & ~ (e3 = op(e3, e4)) & ~ (e2 = op(e3, e4)) & ~ (e1 = op(e3, e4)) & ~ (e0 = op(e3, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP135])])).
fof(f2841, plain, (sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2840, f1585])).
fof(f1585, plain, ~ sP136, inference(subsumption_resolution, [], [f1214, f1076])).
fof(f1076, plain, (e4 = op(e4, unit)), inference(definition_unfolding, [], [f420, f425])).
fof(f420, plain, (e4 = op(e4, e0)), inference(cnf_transformation, [], [f2])).
fof(f1214, plain, (~ (e4 = op(e4, unit)) | ~ sP136), inference(definition_unfolding, [], [f700, f425])).
fof(f700, plain, (~ (e4 = op(e4, e0)) | ~ sP136), inference(cnf_transformation, [], [f253])).
fof(f253, plain, ((~ (e4 = op(e4, e0)) & ~ (e3 = op(e4, e0)) & ~ (e2 = op(e4, e0)) & ~ (e1 = op(e4, e0)) & ~ (e0 = op(e4, e0))) | ~ sP136), inference(nnf_transformation, [], [f143])).
fof(f143, plain, ((~ (e4 = op(e4, e0)) & ~ (e3 = op(e4, e0)) & ~ (e2 = op(e4, e0)) & ~ (e1 = op(e4, e0)) & ~ (e0 = op(e4, e0))) | ~ sP136), inference(usedef, [], [e143])).
fof(e143, plain, (sP136 <=> (~ (e4 = op(e4, e0)) & ~ (e3 = op(e4, e0)) & ~ (e2 = op(e4, e0)) & ~ (e1 = op(e4, e0)) & ~ (e0 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP136])])).
fof(f2840, plain, (sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2839, f1584])).
fof(f1584, plain, ~ sP137, inference(subsumption_resolution, [], [f694, f421])).
fof(f421, plain, (e3 = op(e4, e1)), inference(cnf_transformation, [], [f2])).
fof(f694, plain, (~ (e3 = op(e4, e1)) | ~ sP137), inference(cnf_transformation, [], [f252])).
fof(f252, plain, ((~ (e4 = op(e4, e1)) & ~ (e3 = op(e4, e1)) & ~ (e2 = op(e4, e1)) & ~ (e1 = op(e4, e1)) & ~ (e0 = op(e4, e1))) | ~ sP137), inference(nnf_transformation, [], [f144])).
fof(f144, plain, ((~ (e4 = op(e4, e1)) & ~ (e3 = op(e4, e1)) & ~ (e2 = op(e4, e1)) & ~ (e1 = op(e4, e1)) & ~ (e0 = op(e4, e1))) | ~ sP137), inference(usedef, [], [e144])).
fof(e144, plain, (sP137 <=> (~ (e4 = op(e4, e1)) & ~ (e3 = op(e4, e1)) & ~ (e2 = op(e4, e1)) & ~ (e1 = op(e4, e1)) & ~ (e0 = op(e4, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP137])])).
fof(f2839, plain, (sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2838, f1582])).
fof(f1582, plain, ~ sP138, inference(subsumption_resolution, [], [f1212, f1075])).
fof(f1075, plain, (op(e4, e2) = unit), inference(definition_unfolding, [], [f422, f425])).
fof(f422, plain, (e0 = op(e4, e2)), inference(cnf_transformation, [], [f2])).
fof(f1212, plain, (~ (op(e4, e2) = unit) | ~ sP138), inference(definition_unfolding, [], [f686, f425])).
fof(f686, plain, (~ (e0 = op(e4, e2)) | ~ sP138), inference(cnf_transformation, [], [f251])).
fof(f251, plain, ((~ (e4 = op(e4, e2)) & ~ (e3 = op(e4, e2)) & ~ (e2 = op(e4, e2)) & ~ (e1 = op(e4, e2)) & ~ (e0 = op(e4, e2))) | ~ sP138), inference(nnf_transformation, [], [f145])).
fof(f145, plain, ((~ (e4 = op(e4, e2)) & ~ (e3 = op(e4, e2)) & ~ (e2 = op(e4, e2)) & ~ (e1 = op(e4, e2)) & ~ (e0 = op(e4, e2))) | ~ sP138), inference(usedef, [], [e145])).
fof(e145, plain, (sP138 <=> (~ (e4 = op(e4, e2)) & ~ (e3 = op(e4, e2)) & ~ (e2 = op(e4, e2)) & ~ (e1 = op(e4, e2)) & ~ (e0 = op(e4, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP138])])).
fof(f2838, plain, (sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2837, f1577])).
fof(f1577, plain, ~ sP139, inference(subsumption_resolution, [], [f682, f423])).
fof(f423, plain, (e1 = op(e4, e3)), inference(cnf_transformation, [], [f2])).
fof(f682, plain, (~ (e1 = op(e4, e3)) | ~ sP139), inference(cnf_transformation, [], [f250])).
fof(f250, plain, ((~ (e4 = op(e4, e3)) & ~ (e3 = op(e4, e3)) & ~ (e2 = op(e4, e3)) & ~ (e1 = op(e4, e3)) & ~ (e0 = op(e4, e3))) | ~ sP139), inference(nnf_transformation, [], [f146])).
fof(f146, plain, ((~ (e4 = op(e4, e3)) & ~ (e3 = op(e4, e3)) & ~ (e2 = op(e4, e3)) & ~ (e1 = op(e4, e3)) & ~ (e0 = op(e4, e3))) | ~ sP139), inference(usedef, [], [e146])).
fof(e146, plain, (sP139 <=> (~ (e4 = op(e4, e3)) & ~ (e3 = op(e4, e3)) & ~ (e2 = op(e4, e3)) & ~ (e1 = op(e4, e3)) & ~ (e0 = op(e4, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP139])])).
fof(f2837, plain, (sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2836, f1573])).
fof(f1573, plain, ~ sP140, inference(subsumption_resolution, [], [f678, f424])).
fof(f424, plain, (e2 = op(e4, e4)), inference(cnf_transformation, [], [f2])).
fof(f678, plain, (~ (e2 = op(e4, e4)) | ~ sP140), inference(cnf_transformation, [], [f249])).
fof(f249, plain, ((~ (e4 = op(e4, e4)) & ~ (e3 = op(e4, e4)) & ~ (e2 = op(e4, e4)) & ~ (e1 = op(e4, e4)) & ~ (e0 = op(e4, e4))) | ~ sP140), inference(nnf_transformation, [], [f147])).
fof(f147, plain, ((~ (e4 = op(e4, e4)) & ~ (e3 = op(e4, e4)) & ~ (e2 = op(e4, e4)) & ~ (e1 = op(e4, e4)) & ~ (e0 = op(e4, e4))) | ~ sP140), inference(usedef, [], [e147])).
fof(e147, plain, (sP140 <=> (~ (e4 = op(e4, e4)) & ~ (e3 = op(e4, e4)) & ~ (e2 = op(e4, e4)) & ~ (e1 = op(e4, e4)) & ~ (e0 = op(e4, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP140])])).
fof(f2836, plain, (sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2835, f1087])).
fof(f2835, plain, (~ (unit = op(unit, unit)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2834, f1086])).
fof(f2834, plain, (~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2833, f1082])).
fof(f2833, plain, (~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2832, f1085])).
fof(f2832, plain, (~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2831, f1080])).
fof(f2831, plain, (~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2830, f1084])).
fof(f2830, plain, (~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2829, f1078])).
fof(f2829, plain, (~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2828, f1076])).
fof(f2828, plain, (~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2827, f1425])).
fof(f1425, plain, ~ sP141, inference(trivial_inequality_removal, [], [f1209])).
fof(f1209, plain, (~ (unit = unit) | ~ sP141), inference(definition_unfolding, [], [f671, f425])).
fof(f671, plain, (~ (e0 = unit) | ~ sP141), inference(cnf_transformation, [], [f248])).
fof(f248, plain, ((~ (e4 = unit) & ~ (e3 = unit) & ~ (e2 = unit) & ~ (e1 = unit) & ~ (e0 = unit)) | ~ sP141), inference(nnf_transformation, [], [f148])).
fof(f148, plain, ((~ (e4 = unit) & ~ (e3 = unit) & ~ (e2 = unit) & ~ (e1 = unit) & ~ (e0 = unit)) | ~ sP141), inference(usedef, [], [e148])).
fof(e148, plain, (sP141 <=> (~ (e4 = unit) & ~ (e3 = unit) & ~ (e2 = unit) & ~ (e1 = unit) & ~ (e0 = unit))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP141])])).
fof(f2827, plain, (sP141 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2826, f1570])).
fof(f1570, plain, ~ sP142, inference(subsumption_resolution, [], [f1208, f1087])).
fof(f1208, plain, (~ (unit = op(unit, unit)) | ~ sP142), inference(definition_unfolding, [], [f666, f425, f425, f425])).
fof(f666, plain, (~ (e0 = op(e0, e0)) | ~ sP142), inference(cnf_transformation, [], [f247])).
fof(f247, plain, ((~ (e0 = op(e0, e4)) & ~ (e0 = op(e0, e3)) & ~ (e0 = op(e0, e2)) & ~ (e0 = op(e0, e1)) & ~ (e0 = op(e0, e0))) | ~ sP142), inference(nnf_transformation, [], [f149])).
fof(f149, plain, ((~ (e0 = op(e0, e4)) & ~ (e0 = op(e0, e3)) & ~ (e0 = op(e0, e2)) & ~ (e0 = op(e0, e1)) & ~ (e0 = op(e0, e0))) | ~ sP142), inference(usedef, [], [e149])).
fof(e149, plain, (sP142 <=> (~ (e0 = op(e0, e4)) & ~ (e0 = op(e0, e3)) & ~ (e0 = op(e0, e2)) & ~ (e0 = op(e0, e1)) & ~ (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP142])])).
fof(f2826, plain, (sP142 | sP141 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2825, f1565])).
fof(f1565, plain, ~ sP143, inference(subsumption_resolution, [], [f1203, f1087])).
fof(f1203, plain, (~ (unit = op(unit, unit)) | ~ sP143), inference(definition_unfolding, [], [f661, f425, f425, f425])).
fof(f661, plain, (~ (e0 = op(e0, e0)) | ~ sP143), inference(cnf_transformation, [], [f246])).
fof(f246, plain, ((~ (e0 = op(e4, e0)) & ~ (e0 = op(e3, e0)) & ~ (e0 = op(e2, e0)) & ~ (e0 = op(e1, e0)) & ~ (e0 = op(e0, e0))) | ~ sP143), inference(nnf_transformation, [], [f150])).
fof(f150, plain, ((~ (e0 = op(e4, e0)) & ~ (e0 = op(e3, e0)) & ~ (e0 = op(e2, e0)) & ~ (e0 = op(e1, e0)) & ~ (e0 = op(e0, e0))) | ~ sP143), inference(usedef, [], [e150])).
fof(e150, plain, (sP143 <=> (~ (e0 = op(e4, e0)) & ~ (e0 = op(e3, e0)) & ~ (e0 = op(e2, e0)) & ~ (e0 = op(e1, e0)) & ~ (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP143])])).
fof(f2825, plain, (sP143 | sP142 | sP141 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2824, f1560])).
fof(f1560, plain, ~ sP144, inference(subsumption_resolution, [], [f1197, f1086])).
fof(f1197, plain, (~ (e1 = op(unit, e1)) | ~ sP144), inference(definition_unfolding, [], [f657, f425])).
fof(f657, plain, (~ (e1 = op(e0, e1)) | ~ sP144), inference(cnf_transformation, [], [f245])).
fof(f245, plain, ((~ (e1 = op(e0, e4)) & ~ (e1 = op(e0, e3)) & ~ (e1 = op(e0, e2)) & ~ (e1 = op(e0, e1)) & ~ (e1 = op(e0, e0))) | ~ sP144), inference(nnf_transformation, [], [f151])).
fof(f151, plain, ((~ (e1 = op(e0, e4)) & ~ (e1 = op(e0, e3)) & ~ (e1 = op(e0, e2)) & ~ (e1 = op(e0, e1)) & ~ (e1 = op(e0, e0))) | ~ sP144), inference(usedef, [], [e151])).
fof(e151, plain, (sP144 <=> (~ (e1 = op(e0, e4)) & ~ (e1 = op(e0, e3)) & ~ (e1 = op(e0, e2)) & ~ (e1 = op(e0, e1)) & ~ (e1 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP144])])).
fof(f2824, plain, (sP144 | sP143 | sP142 | sP141 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2823, f1556])).
fof(f1556, plain, ~ sP145, inference(subsumption_resolution, [], [f1192, f1082])).
fof(f1192, plain, (~ (e1 = op(e1, unit)) | ~ sP145), inference(definition_unfolding, [], [f652, f425])).
fof(f652, plain, (~ (e1 = op(e1, e0)) | ~ sP145), inference(cnf_transformation, [], [f244])).
fof(f244, plain, ((~ (e1 = op(e4, e0)) & ~ (e1 = op(e3, e0)) & ~ (e1 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e1 = op(e0, e0))) | ~ sP145), inference(nnf_transformation, [], [f152])).
fof(f152, plain, ((~ (e1 = op(e4, e0)) & ~ (e1 = op(e3, e0)) & ~ (e1 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e1 = op(e0, e0))) | ~ sP145), inference(usedef, [], [e152])).
fof(e152, plain, (sP145 <=> (~ (e1 = op(e4, e0)) & ~ (e1 = op(e3, e0)) & ~ (e1 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e1 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP145])])).
fof(f2823, plain, (sP145 | sP144 | sP143 | sP142 | sP141 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2822, f1552])).
fof(f1552, plain, ~ sP146, inference(subsumption_resolution, [], [f1186, f1085])).
fof(f1186, plain, (~ (e2 = op(unit, e2)) | ~ sP146), inference(definition_unfolding, [], [f648, f425])).
fof(f648, plain, (~ (e2 = op(e0, e2)) | ~ sP146), inference(cnf_transformation, [], [f243])).
fof(f243, plain, ((~ (e2 = op(e0, e4)) & ~ (e2 = op(e0, e3)) & ~ (e2 = op(e0, e2)) & ~ (e2 = op(e0, e1)) & ~ (e2 = op(e0, e0))) | ~ sP146), inference(nnf_transformation, [], [f153])).
fof(f153, plain, ((~ (e2 = op(e0, e4)) & ~ (e2 = op(e0, e3)) & ~ (e2 = op(e0, e2)) & ~ (e2 = op(e0, e1)) & ~ (e2 = op(e0, e0))) | ~ sP146), inference(usedef, [], [e153])).
fof(e153, plain, (sP146 <=> (~ (e2 = op(e0, e4)) & ~ (e2 = op(e0, e3)) & ~ (e2 = op(e0, e2)) & ~ (e2 = op(e0, e1)) & ~ (e2 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP146])])).
fof(f2822, plain, (sP146 | sP145 | sP144 | sP143 | sP142 | sP141 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2821, f1549])).
fof(f1549, plain, ~ sP147, inference(subsumption_resolution, [], [f1181, f1080])).
fof(f1181, plain, (~ (e2 = op(e2, unit)) | ~ sP147), inference(definition_unfolding, [], [f643, f425])).
fof(f643, plain, (~ (e2 = op(e2, e0)) | ~ sP147), inference(cnf_transformation, [], [f242])).
fof(f242, plain, ((~ (e2 = op(e4, e0)) & ~ (e2 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e2 = op(e1, e0)) & ~ (e2 = op(e0, e0))) | ~ sP147), inference(nnf_transformation, [], [f154])).
fof(f154, plain, ((~ (e2 = op(e4, e0)) & ~ (e2 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e2 = op(e1, e0)) & ~ (e2 = op(e0, e0))) | ~ sP147), inference(usedef, [], [e154])).
fof(e154, plain, (sP147 <=> (~ (e2 = op(e4, e0)) & ~ (e2 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e2 = op(e1, e0)) & ~ (e2 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP147])])).
fof(f2821, plain, (sP147 | sP146 | sP145 | sP144 | sP143 | sP142 | sP141 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2820, f1546])).
fof(f1546, plain, ~ sP148, inference(subsumption_resolution, [], [f1175, f1084])).
fof(f1175, plain, (~ (e3 = op(unit, e3)) | ~ sP148), inference(definition_unfolding, [], [f639, f425])).
fof(f639, plain, (~ (e3 = op(e0, e3)) | ~ sP148), inference(cnf_transformation, [], [f241])).
fof(f241, plain, ((~ (e3 = op(e0, e4)) & ~ (e3 = op(e0, e3)) & ~ (e3 = op(e0, e2)) & ~ (e3 = op(e0, e1)) & ~ (e3 = op(e0, e0))) | ~ sP148), inference(nnf_transformation, [], [f155])).
fof(f155, plain, ((~ (e3 = op(e0, e4)) & ~ (e3 = op(e0, e3)) & ~ (e3 = op(e0, e2)) & ~ (e3 = op(e0, e1)) & ~ (e3 = op(e0, e0))) | ~ sP148), inference(usedef, [], [e155])).
fof(e155, plain, (sP148 <=> (~ (e3 = op(e0, e4)) & ~ (e3 = op(e0, e3)) & ~ (e3 = op(e0, e2)) & ~ (e3 = op(e0, e1)) & ~ (e3 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP148])])).
fof(f2820, plain, (sP148 | sP147 | sP146 | sP145 | sP144 | sP143 | sP142 | sP141 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2819, f1544])).
fof(f1544, plain, ~ sP149, inference(subsumption_resolution, [], [f1170, f1078])).
fof(f1170, plain, (~ (e3 = op(e3, unit)) | ~ sP149), inference(definition_unfolding, [], [f634, f425])).
fof(f634, plain, (~ (e3 = op(e3, e0)) | ~ sP149), inference(cnf_transformation, [], [f240])).
fof(f240, plain, ((~ (e3 = op(e4, e0)) & ~ (e3 = op(e3, e0)) & ~ (e3 = op(e2, e0)) & ~ (e3 = op(e1, e0)) & ~ (e3 = op(e0, e0))) | ~ sP149), inference(nnf_transformation, [], [f156])).
fof(f156, plain, ((~ (e3 = op(e4, e0)) & ~ (e3 = op(e3, e0)) & ~ (e3 = op(e2, e0)) & ~ (e3 = op(e1, e0)) & ~ (e3 = op(e0, e0))) | ~ sP149), inference(usedef, [], [e156])).
fof(e156, plain, (sP149 <=> (~ (e3 = op(e4, e0)) & ~ (e3 = op(e3, e0)) & ~ (e3 = op(e2, e0)) & ~ (e3 = op(e1, e0)) & ~ (e3 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP149])])).
fof(f2819, plain, (sP149 | sP148 | sP147 | sP146 | sP145 | sP144 | sP143 | sP142 | sP141 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2818, f1542])).
fof(f1542, plain, ~ sP150, inference(subsumption_resolution, [], [f1164, f1083])).
fof(f1164, plain, (~ (e4 = op(unit, e4)) | ~ sP150), inference(definition_unfolding, [], [f630, f425])).
fof(f630, plain, (~ (e4 = op(e0, e4)) | ~ sP150), inference(cnf_transformation, [], [f239])).
fof(f239, plain, ((~ (e4 = op(e0, e4)) & ~ (e4 = op(e0, e3)) & ~ (e4 = op(e0, e2)) & ~ (e4 = op(e0, e1)) & ~ (e4 = op(e0, e0))) | ~ sP150), inference(nnf_transformation, [], [f157])).
fof(f157, plain, ((~ (e4 = op(e0, e4)) & ~ (e4 = op(e0, e3)) & ~ (e4 = op(e0, e2)) & ~ (e4 = op(e0, e1)) & ~ (e4 = op(e0, e0))) | ~ sP150), inference(usedef, [], [e157])).
fof(e157, plain, (sP150 <=> (~ (e4 = op(e0, e4)) & ~ (e4 = op(e0, e3)) & ~ (e4 = op(e0, e2)) & ~ (e4 = op(e0, e1)) & ~ (e4 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP150])])).
fof(f2818, plain, (sP150 | sP149 | sP148 | sP147 | sP146 | sP145 | sP144 | sP143 | sP142 | sP141 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2817, f1541])).
fof(f1541, plain, ~ sP151, inference(subsumption_resolution, [], [f1159, f1076])).
fof(f1159, plain, (~ (e4 = op(e4, unit)) | ~ sP151), inference(definition_unfolding, [], [f625, f425])).
fof(f625, plain, (~ (e4 = op(e4, e0)) | ~ sP151), inference(cnf_transformation, [], [f238])).
fof(f238, plain, ((~ (e4 = op(e4, e0)) & ~ (e4 = op(e3, e0)) & ~ (e4 = op(e2, e0)) & ~ (e4 = op(e1, e0)) & ~ (e4 = op(e0, e0))) | ~ sP151), inference(nnf_transformation, [], [f158])).
fof(f158, plain, ((~ (e4 = op(e4, e0)) & ~ (e4 = op(e3, e0)) & ~ (e4 = op(e2, e0)) & ~ (e4 = op(e1, e0)) & ~ (e4 = op(e0, e0))) | ~ sP151), inference(usedef, [], [e158])).
fof(e158, plain, (sP151 <=> (~ (e4 = op(e4, e0)) & ~ (e4 = op(e3, e0)) & ~ (e4 = op(e2, e0)) & ~ (e4 = op(e1, e0)) & ~ (e4 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP151])])).
fof(f2817, plain, (sP151 | sP150 | sP149 | sP148 | sP147 | sP146 | sP145 | sP144 | sP143 | sP142 | sP141 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2816, f1540])).
fof(f1540, plain, ~ sP152, inference(subsumption_resolution, [], [f1157, f1081])).
fof(f1157, plain, (~ (op(e1, e1) = unit) | ~ sP152), inference(definition_unfolding, [], [f617, f425])).
fof(f617, plain, (~ (e0 = op(e1, e1)) | ~ sP152), inference(cnf_transformation, [], [f237])).
fof(f237, plain, ((~ (e0 = op(e1, e4)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e1, e0))) | ~ sP152), inference(nnf_transformation, [], [f159])).
fof(f159, plain, ((~ (e0 = op(e1, e4)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e1, e0))) | ~ sP152), inference(usedef, [], [e159])).
fof(e159, plain, (sP152 <=> (~ (e0 = op(e1, e4)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP152])])).
fof(f2816, plain, (sP152 | sP151 | sP150 | sP149 | sP148 | sP147 | sP146 | sP145 | sP144 | sP143 | sP142 | sP141 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2815, f1536])).
fof(f1536, plain, ~ sP153, inference(subsumption_resolution, [], [f1152, f1081])).
fof(f1152, plain, (~ (op(e1, e1) = unit) | ~ sP153), inference(definition_unfolding, [], [f612, f425])).
fof(f612, plain, (~ (e0 = op(e1, e1)) | ~ sP153), inference(cnf_transformation, [], [f236])).
fof(f236, plain, ((~ (e0 = op(e4, e1)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e0, e1))) | ~ sP153), inference(nnf_transformation, [], [f160])).
fof(f160, plain, ((~ (e0 = op(e4, e1)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e0, e1))) | ~ sP153), inference(usedef, [], [e160])).
fof(e160, plain, (sP153 <=> (~ (e0 = op(e4, e1)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP153])])).
fof(f2815, plain, (sP153 | sP152 | sP151 | sP150 | sP149 | sP148 | sP147 | sP146 | sP145 | sP144 | sP143 | sP142 | sP141 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2814, f1532])).
fof(f1532, plain, ~ sP154, inference(subsumption_resolution, [], [f1148, f1082])).
fof(f1148, plain, (~ (e1 = op(e1, unit)) | ~ sP154), inference(definition_unfolding, [], [f606, f425])).
fof(f606, plain, (~ (e1 = op(e1, e0)) | ~ sP154), inference(cnf_transformation, [], [f235])).
fof(f235, plain, ((~ (e1 = op(e1, e4)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e1, e0))) | ~ sP154), inference(nnf_transformation, [], [f161])).
fof(f161, plain, ((~ (e1 = op(e1, e4)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e1, e0))) | ~ sP154), inference(usedef, [], [e161])).
fof(e161, plain, (sP154 <=> (~ (e1 = op(e1, e4)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP154])])).
fof(f2814, plain, (sP154 | sP153 | sP152 | sP151 | sP150 | sP149 | sP148 | sP147 | sP146 | sP145 | sP144 | sP143 | sP142 | sP141 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2813, f1527])).
fof(f1527, plain, ~ sP155, inference(subsumption_resolution, [], [f1147, f1086])).
fof(f1147, plain, (~ (e1 = op(unit, e1)) | ~ sP155), inference(definition_unfolding, [], [f601, f425])).
fof(f601, plain, (~ (e1 = op(e0, e1)) | ~ sP155), inference(cnf_transformation, [], [f234])).
fof(f234, plain, ((~ (e1 = op(e4, e1)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e0, e1))) | ~ sP155), inference(nnf_transformation, [], [f162])).
fof(f162, plain, ((~ (e1 = op(e4, e1)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e0, e1))) | ~ sP155), inference(usedef, [], [e162])).
fof(e162, plain, (sP155 <=> (~ (e1 = op(e4, e1)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP155])])).
fof(f2813, plain, (sP155 | sP154 | sP153 | sP152 | sP151 | sP150 | sP149 | sP148 | sP147 | sP146 | sP145 | sP144 | sP143 | sP142 | sP141 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2812, f1522])).
fof(f1522, plain, ~ sP156, inference(subsumption_resolution, [], [f599, f408])).
fof(f599, plain, (~ (e2 = op(e1, e3)) | ~ sP156), inference(cnf_transformation, [], [f233])).
fof(f233, plain, ((~ (e2 = op(e1, e4)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e1, e0))) | ~ sP156), inference(nnf_transformation, [], [f163])).
fof(f163, plain, ((~ (e2 = op(e1, e4)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e1, e0))) | ~ sP156), inference(usedef, [], [e163])).
fof(e163, plain, (sP156 <=> (~ (e2 = op(e1, e4)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP156])])).
fof(f2812, plain, (sP156 | sP155 | sP154 | sP153 | sP152 | sP151 | sP150 | sP149 | sP148 | sP147 | sP146 | sP145 | sP144 | sP143 | sP142 | sP141 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2811, f1520])).
fof(f1520, plain, ~ sP157, inference(subsumption_resolution, [], [f594, f416])).
fof(f594, plain, (~ (e2 = op(e3, e1)) | ~ sP157), inference(cnf_transformation, [], [f232])).
fof(f232, plain, ((~ (e2 = op(e4, e1)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e0, e1))) | ~ sP157), inference(nnf_transformation, [], [f164])).
fof(f164, plain, ((~ (e2 = op(e4, e1)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e0, e1))) | ~ sP157), inference(usedef, [], [e164])).
fof(e164, plain, (sP157 <=> (~ (e2 = op(e4, e1)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP157])])).
fof(f2811, plain, (sP157 | sP156 | sP155 | sP154 | sP153 | sP152 | sP151 | sP150 | sP149 | sP148 | sP147 | sP146 | sP145 | sP144 | sP143 | sP142 | sP141 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2810, f1518])).
fof(f1518, plain, ~ sP158, inference(subsumption_resolution, [], [f590, f409])).
fof(f590, plain, (~ (e3 = op(e1, e4)) | ~ sP158), inference(cnf_transformation, [], [f231])).
fof(f231, plain, ((~ (e3 = op(e1, e4)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e1, e0))) | ~ sP158), inference(nnf_transformation, [], [f165])).
fof(f165, plain, ((~ (e3 = op(e1, e4)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e1, e0))) | ~ sP158), inference(usedef, [], [e165])).
fof(e165, plain, (sP158 <=> (~ (e3 = op(e1, e4)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP158])])).
fof(f2810, plain, (sP158 | sP157 | sP156 | sP155 | sP154 | sP153 | sP152 | sP151 | sP150 | sP149 | sP148 | sP147 | sP146 | sP145 | sP144 | sP143 | sP142 | sP141 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2809, f1517])).
fof(f1517, plain, ~ sP159, inference(subsumption_resolution, [], [f585, f421])).
fof(f585, plain, (~ (e3 = op(e4, e1)) | ~ sP159), inference(cnf_transformation, [], [f230])).
fof(f230, plain, ((~ (e3 = op(e4, e1)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e0, e1))) | ~ sP159), inference(nnf_transformation, [], [f166])).
fof(f166, plain, ((~ (e3 = op(e4, e1)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e0, e1))) | ~ sP159), inference(usedef, [], [e166])).
fof(e166, plain, (sP159 <=> (~ (e3 = op(e4, e1)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP159])])).
fof(f2809, plain, (sP159 | sP158 | sP157 | sP156 | sP155 | sP154 | sP153 | sP152 | sP151 | sP150 | sP149 | sP148 | sP147 | sP146 | sP145 | sP144 | sP143 | sP142 | sP141 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2808, f1516])).
fof(f1516, plain, ~ sP160, inference(subsumption_resolution, [], [f578, f407])).
fof(f578, plain, (~ (e4 = op(e1, e2)) | ~ sP160), inference(cnf_transformation, [], [f229])).
fof(f229, plain, ((~ (e4 = op(e1, e4)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e1, e0))) | ~ sP160), inference(nnf_transformation, [], [f167])).
fof(f167, plain, ((~ (e4 = op(e1, e4)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e1, e0))) | ~ sP160), inference(usedef, [], [e167])).
fof(e167, plain, (sP160 <=> (~ (e4 = op(e1, e4)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP160])])).
fof(f2808, plain, (sP160 | sP159 | sP158 | sP157 | sP156 | sP155 | sP154 | sP153 | sP152 | sP151 | sP150 | sP149 | sP148 | sP147 | sP146 | sP145 | sP144 | sP143 | sP142 | sP141 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2807, f1513])).
fof(f1513, plain, ~ sP161, inference(subsumption_resolution, [], [f573, f411])).
fof(f573, plain, (~ (e4 = op(e2, e1)) | ~ sP161), inference(cnf_transformation, [], [f228])).
fof(f228, plain, ((~ (e4 = op(e4, e1)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e0, e1))) | ~ sP161), inference(nnf_transformation, [], [f168])).
fof(f168, plain, ((~ (e4 = op(e4, e1)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e0, e1))) | ~ sP161), inference(usedef, [], [e168])).
fof(e168, plain, (sP161 <=> (~ (e4 = op(e4, e1)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP161])])).
fof(f2807, plain, (sP161 | sP160 | sP159 | sP158 | sP157 | sP156 | sP155 | sP154 | sP153 | sP152 | sP151 | sP150 | sP149 | sP148 | sP147 | sP146 | sP145 | sP144 | sP143 | sP142 | sP141 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2806, f1510])).
fof(f1510, plain, ~ sP162, inference(subsumption_resolution, [], [f1137, f1079])).
fof(f1137, plain, (~ (op(e2, e3) = unit) | ~ sP162), inference(definition_unfolding, [], [f569, f425])).
fof(f569, plain, (~ (e0 = op(e2, e3)) | ~ sP162), inference(cnf_transformation, [], [f227])).
fof(f227, plain, ((~ (e0 = op(e2, e4)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e2, e0))) | ~ sP162), inference(nnf_transformation, [], [f169])).
fof(f169, plain, ((~ (e0 = op(e2, e4)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e2, e0))) | ~ sP162), inference(usedef, [], [e169])).
fof(e169, plain, (sP162 <=> (~ (e0 = op(e2, e4)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP162])])).
fof(f2806, plain, (sP162 | sP161 | sP160 | sP159 | sP158 | sP157 | sP156 | sP155 | sP154 | sP153 | sP152 | sP151 | sP150 | sP149 | sP148 | sP147 | sP146 | sP145 | sP144 | sP143 | sP142 | sP141 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2805, f1508])).
fof(f1508, plain, ~ sP163, inference(subsumption_resolution, [], [f1131, f1075])).
fof(f1131, plain, (~ (op(e4, e2) = unit) | ~ sP163), inference(definition_unfolding, [], [f565, f425])).
fof(f565, plain, (~ (e0 = op(e4, e2)) | ~ sP163), inference(cnf_transformation, [], [f226])).
fof(f226, plain, ((~ (e0 = op(e4, e2)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e0, e2))) | ~ sP163), inference(nnf_transformation, [], [f170])).
fof(f170, plain, ((~ (e0 = op(e4, e2)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e0, e2))) | ~ sP163), inference(usedef, [], [e170])).
fof(e170, plain, (sP163 <=> (~ (e0 = op(e4, e2)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP163])])).
fof(f2805, plain, (sP163 | sP162 | sP161 | sP160 | sP159 | sP158 | sP157 | sP156 | sP155 | sP154 | sP153 | sP152 | sP151 | sP150 | sP149 | sP148 | sP147 | sP146 | sP145 | sP144 | sP143 | sP142 | sP141 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2804, f1507])).
fof(f1507, plain, ~ sP164, inference(subsumption_resolution, [], [f560, f414])).
fof(f560, plain, (~ (e1 = op(e2, e4)) | ~ sP164), inference(cnf_transformation, [], [f225])).
fof(f225, plain, ((~ (e1 = op(e2, e4)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e2, e0))) | ~ sP164), inference(nnf_transformation, [], [f171])).
fof(f171, plain, ((~ (e1 = op(e2, e4)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e2, e0))) | ~ sP164), inference(usedef, [], [e171])).
fof(e171, plain, (sP164 <=> (~ (e1 = op(e2, e4)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP164])])).
fof(f2804, plain, (sP164 | sP163 | sP162 | sP161 | sP160 | sP159 | sP158 | sP157 | sP156 | sP155 | sP154 | sP153 | sP152 | sP151 | sP150 | sP149 | sP148 | sP147 | sP146 | sP145 | sP144 | sP143 | sP142 | sP141 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2803, f1506])).
fof(f1506, plain, ~ sP165, inference(subsumption_resolution, [], [f554, f417])).
fof(f554, plain, (~ (e1 = op(e3, e2)) | ~ sP165), inference(cnf_transformation, [], [f224])).
fof(f224, plain, ((~ (e1 = op(e4, e2)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e0, e2))) | ~ sP165), inference(nnf_transformation, [], [f172])).
fof(f172, plain, ((~ (e1 = op(e4, e2)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e0, e2))) | ~ sP165), inference(usedef, [], [e172])).
fof(e172, plain, (sP165 <=> (~ (e1 = op(e4, e2)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP165])])).
fof(f2803, plain, (sP165 | sP164 | sP163 | sP162 | sP161 | sP160 | sP159 | sP158 | sP157 | sP156 | sP155 | sP154 | sP153 | sP152 | sP151 | sP150 | sP149 | sP148 | sP147 | sP146 | sP145 | sP144 | sP143 | sP142 | sP141 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2802, f1504])).
fof(f1504, plain, ~ sP166, inference(subsumption_resolution, [], [f1128, f1080])).
fof(f1128, plain, (~ (e2 = op(e2, unit)) | ~ sP166), inference(definition_unfolding, [], [f546, f425])).
fof(f546, plain, (~ (e2 = op(e2, e0)) | ~ sP166), inference(cnf_transformation, [], [f223])).
fof(f223, plain, ((~ (e2 = op(e2, e4)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e2, e0))) | ~ sP166), inference(nnf_transformation, [], [f173])).
fof(f173, plain, ((~ (e2 = op(e2, e4)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e2, e0))) | ~ sP166), inference(usedef, [], [e173])).
fof(e173, plain, (sP166 <=> (~ (e2 = op(e2, e4)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP166])])).
fof(f2802, plain, (sP166 | sP165 | sP164 | sP163 | sP162 | sP161 | sP160 | sP159 | sP158 | sP157 | sP156 | sP155 | sP154 | sP153 | sP152 | sP151 | sP150 | sP149 | sP148 | sP147 | sP146 | sP145 | sP144 | sP143 | sP142 | sP141 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2801, f1499])).
fof(f1499, plain, ~ sP167, inference(subsumption_resolution, [], [f1127, f1085])).
fof(f1127, plain, (~ (e2 = op(unit, e2)) | ~ sP167), inference(definition_unfolding, [], [f541, f425])).
fof(f541, plain, (~ (e2 = op(e0, e2)) | ~ sP167), inference(cnf_transformation, [], [f222])).
fof(f222, plain, ((~ (e2 = op(e4, e2)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e0, e2))) | ~ sP167), inference(nnf_transformation, [], [f174])).
fof(f174, plain, ((~ (e2 = op(e4, e2)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e0, e2))) | ~ sP167), inference(usedef, [], [e174])).
fof(e174, plain, (sP167 <=> (~ (e2 = op(e4, e2)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP167])])).
fof(f2801, plain, (sP167 | sP166 | sP165 | sP164 | sP163 | sP162 | sP161 | sP160 | sP159 | sP158 | sP157 | sP156 | sP155 | sP154 | sP153 | sP152 | sP151 | sP150 | sP149 | sP148 | sP147 | sP146 | sP145 | sP144 | sP143 | sP142 | sP141 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2800, f1494])).
fof(f1494, plain, ~ sP168, inference(subsumption_resolution, [], [f538, f412])).
fof(f538, plain, (~ (e3 = op(e2, e2)) | ~ sP168), inference(cnf_transformation, [], [f221])).
fof(f221, plain, ((~ (e3 = op(e2, e4)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e2, e0))) | ~ sP168), inference(nnf_transformation, [], [f175])).
fof(f175, plain, ((~ (e3 = op(e2, e4)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e2, e0))) | ~ sP168), inference(usedef, [], [e175])).
fof(e175, plain, (sP168 <=> (~ (e3 = op(e2, e4)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP168])])).
fof(f2800, plain, (sP168 | sP167 | sP166 | sP165 | sP164 | sP163 | sP162 | sP161 | sP160 | sP159 | sP158 | sP157 | sP156 | sP155 | sP154 | sP153 | sP152 | sP151 | sP150 | sP149 | sP148 | sP147 | sP146 | sP145 | sP144 | sP143 | sP142 | sP141 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2799, f1491])).
fof(f1491, plain, ~ sP169, inference(subsumption_resolution, [], [f533, f412])).
fof(f533, plain, (~ (e3 = op(e2, e2)) | ~ sP169), inference(cnf_transformation, [], [f220])).
fof(f220, plain, ((~ (e3 = op(e4, e2)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e0, e2))) | ~ sP169), inference(nnf_transformation, [], [f176])).
fof(f176, plain, ((~ (e3 = op(e4, e2)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e0, e2))) | ~ sP169), inference(usedef, [], [e176])).
fof(e176, plain, (sP169 <=> (~ (e3 = op(e4, e2)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP169])])).
fof(f2799, plain, (sP169 | sP168 | sP167 | sP166 | sP165 | sP164 | sP163 | sP162 | sP161 | sP160 | sP159 | sP158 | sP157 | sP156 | sP155 | sP154 | sP153 | sP152 | sP151 | sP150 | sP149 | sP148 | sP147 | sP146 | sP145 | sP144 | sP143 | sP142 | sP141 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2798, f1488])).
fof(f1488, plain, ~ sP170, inference(subsumption_resolution, [], [f527, f411])).
fof(f527, plain, (~ (e4 = op(e2, e1)) | ~ sP170), inference(cnf_transformation, [], [f219])).
fof(f219, plain, ((~ (e4 = op(e2, e4)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e2, e0))) | ~ sP170), inference(nnf_transformation, [], [f177])).
fof(f177, plain, ((~ (e4 = op(e2, e4)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e2, e0))) | ~ sP170), inference(usedef, [], [e177])).
fof(e177, plain, (sP170 <=> (~ (e4 = op(e2, e4)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP170])])).
fof(f2798, plain, (sP170 | sP169 | sP168 | sP167 | sP166 | sP165 | sP164 | sP163 | sP162 | sP161 | sP160 | sP159 | sP158 | sP157 | sP156 | sP155 | sP154 | sP153 | sP152 | sP151 | sP150 | sP149 | sP148 | sP147 | sP146 | sP145 | sP144 | sP143 | sP142 | sP141 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2797, f1484])).
fof(f1484, plain, ~ sP171, inference(subsumption_resolution, [], [f522, f407])).
fof(f522, plain, (~ (e4 = op(e1, e2)) | ~ sP171), inference(cnf_transformation, [], [f218])).
fof(f218, plain, ((~ (e4 = op(e4, e2)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e0, e2))) | ~ sP171), inference(nnf_transformation, [], [f178])).
fof(f178, plain, ((~ (e4 = op(e4, e2)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e0, e2))) | ~ sP171), inference(usedef, [], [e178])).
fof(e178, plain, (sP171 <=> (~ (e4 = op(e4, e2)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP171])])).
fof(f2797, plain, (sP171 | sP170 | sP169 | sP168 | sP167 | sP166 | sP165 | sP164 | sP163 | sP162 | sP161 | sP160 | sP159 | sP158 | sP157 | sP156 | sP155 | sP154 | sP153 | sP152 | sP151 | sP150 | sP149 | sP148 | sP147 | sP146 | sP145 | sP144 | sP143 | sP142 | sP141 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2796, f1480])).
fof(f1480, plain, ~ sP172, inference(subsumption_resolution, [], [f1118, f1077])).
fof(f1118, plain, (~ (op(e3, e4) = unit) | ~ sP172), inference(definition_unfolding, [], [f520, f425])).
fof(f520, plain, (~ (e0 = op(e3, e4)) | ~ sP172), inference(cnf_transformation, [], [f217])).
fof(f217, plain, ((~ (e0 = op(e3, e4)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e3, e0))) | ~ sP172), inference(nnf_transformation, [], [f179])).
fof(f179, plain, ((~ (e0 = op(e3, e4)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e3, e0))) | ~ sP172), inference(usedef, [], [e179])).
fof(e179, plain, (sP172 <=> (~ (e0 = op(e3, e4)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP172])])).
fof(f2796, plain, (sP172 | sP171 | sP170 | sP169 | sP168 | sP167 | sP166 | sP165 | sP164 | sP163 | sP162 | sP161 | sP160 | sP159 | sP158 | sP157 | sP156 | sP155 | sP154 | sP153 | sP152 | sP151 | sP150 | sP149 | sP148 | sP147 | sP146 | sP145 | sP144 | sP143 | sP142 | sP141 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2795, f1479])).
fof(f1479, plain, ~ sP173, inference(subsumption_resolution, [], [f1115, f1079])).
fof(f1115, plain, (~ (op(e2, e3) = unit) | ~ sP173), inference(definition_unfolding, [], [f513, f425])).
fof(f513, plain, (~ (e0 = op(e2, e3)) | ~ sP173), inference(cnf_transformation, [], [f216])).
fof(f216, plain, ((~ (e0 = op(e4, e3)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e0, e3))) | ~ sP173), inference(nnf_transformation, [], [f180])).
fof(f180, plain, ((~ (e0 = op(e4, e3)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e0, e3))) | ~ sP173), inference(usedef, [], [e180])).
fof(e180, plain, (sP173 <=> (~ (e0 = op(e4, e3)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP173])])).
fof(f2795, plain, (sP173 | sP172 | sP171 | sP170 | sP169 | sP168 | sP167 | sP166 | sP165 | sP164 | sP163 | sP162 | sP161 | sP160 | sP159 | sP158 | sP157 | sP156 | sP155 | sP154 | sP153 | sP152 | sP151 | sP150 | sP149 | sP148 | sP147 | sP146 | sP145 | sP144 | sP143 | sP142 | sP141 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2794, f1476])).
fof(f1476, plain, ~ sP174, inference(subsumption_resolution, [], [f508, f417])).
fof(f508, plain, (~ (e1 = op(e3, e2)) | ~ sP174), inference(cnf_transformation, [], [f215])).
fof(f215, plain, ((~ (e1 = op(e3, e4)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e3, e0))) | ~ sP174), inference(nnf_transformation, [], [f181])).
fof(f181, plain, ((~ (e1 = op(e3, e4)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e3, e0))) | ~ sP174), inference(usedef, [], [e181])).
fof(e181, plain, (sP174 <=> (~ (e1 = op(e3, e4)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP174])])).
fof(f2794, plain, (sP174 | sP173 | sP172 | sP171 | sP170 | sP169 | sP168 | sP167 | sP166 | sP165 | sP164 | sP163 | sP162 | sP161 | sP160 | sP159 | sP158 | sP157 | sP156 | sP155 | sP154 | sP153 | sP152 | sP151 | sP150 | sP149 | sP148 | sP147 | sP146 | sP145 | sP144 | sP143 | sP142 | sP141 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2793, f1473])).
fof(f1473, plain, ~ sP175, inference(subsumption_resolution, [], [f505, f423])).
fof(f505, plain, (~ (e1 = op(e4, e3)) | ~ sP175), inference(cnf_transformation, [], [f214])).
fof(f214, plain, ((~ (e1 = op(e4, e3)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e0, e3))) | ~ sP175), inference(nnf_transformation, [], [f182])).
fof(f182, plain, ((~ (e1 = op(e4, e3)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e0, e3))) | ~ sP175), inference(usedef, [], [e182])).
fof(e182, plain, (sP175 <=> (~ (e1 = op(e4, e3)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP175])])).
fof(f2793, plain, (sP175 | sP174 | sP173 | sP172 | sP171 | sP170 | sP169 | sP168 | sP167 | sP166 | sP165 | sP164 | sP163 | sP162 | sP161 | sP160 | sP159 | sP158 | sP157 | sP156 | sP155 | sP154 | sP153 | sP152 | sP151 | sP150 | sP149 | sP148 | sP147 | sP146 | sP145 | sP144 | sP143 | sP142 | sP141 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2792, f1472])).
fof(f1472, plain, ~ sP176, inference(subsumption_resolution, [], [f497, f416])).
fof(f497, plain, (~ (e2 = op(e3, e1)) | ~ sP176), inference(cnf_transformation, [], [f213])).
fof(f213, plain, ((~ (e2 = op(e3, e4)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e3, e0))) | ~ sP176), inference(nnf_transformation, [], [f183])).
fof(f183, plain, ((~ (e2 = op(e3, e4)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e3, e0))) | ~ sP176), inference(usedef, [], [e183])).
fof(e183, plain, (sP176 <=> (~ (e2 = op(e3, e4)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP176])])).
fof(f2792, plain, (sP176 | sP175 | sP174 | sP173 | sP172 | sP171 | sP170 | sP169 | sP168 | sP167 | sP166 | sP165 | sP164 | sP163 | sP162 | sP161 | sP160 | sP159 | sP158 | sP157 | sP156 | sP155 | sP154 | sP153 | sP152 | sP151 | sP150 | sP149 | sP148 | sP147 | sP146 | sP145 | sP144 | sP143 | sP142 | sP141 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2791, f1468])).
fof(f1468, plain, ~ sP177, inference(subsumption_resolution, [], [f492, f408])).
fof(f492, plain, (~ (e2 = op(e1, e3)) | ~ sP177), inference(cnf_transformation, [], [f212])).
fof(f212, plain, ((~ (e2 = op(e4, e3)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e0, e3))) | ~ sP177), inference(nnf_transformation, [], [f184])).
fof(f184, plain, ((~ (e2 = op(e4, e3)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e0, e3))) | ~ sP177), inference(usedef, [], [e184])).
fof(e184, plain, (sP177 <=> (~ (e2 = op(e4, e3)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP177])])).
fof(f2791, plain, (sP177 | sP176 | sP175 | sP174 | sP173 | sP172 | sP171 | sP170 | sP169 | sP168 | sP167 | sP166 | sP165 | sP164 | sP163 | sP162 | sP161 | sP160 | sP159 | sP158 | sP157 | sP156 | sP155 | sP154 | sP153 | sP152 | sP151 | sP150 | sP149 | sP148 | sP147 | sP146 | sP145 | sP144 | sP143 | sP142 | sP141 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2790, f1464])).
fof(f1464, plain, ~ sP178, inference(subsumption_resolution, [], [f1108, f1078])).
fof(f1108, plain, (~ (e3 = op(e3, unit)) | ~ sP178), inference(definition_unfolding, [], [f486, f425])).
fof(f486, plain, (~ (e3 = op(e3, e0)) | ~ sP178), inference(cnf_transformation, [], [f211])).
fof(f211, plain, ((~ (e3 = op(e3, e4)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e3, e0))) | ~ sP178), inference(nnf_transformation, [], [f185])).
fof(f185, plain, ((~ (e3 = op(e3, e4)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e3, e0))) | ~ sP178), inference(usedef, [], [e185])).
fof(e185, plain, (sP178 <=> (~ (e3 = op(e3, e4)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP178])])).
fof(f2790, plain, (sP178 | sP177 | sP176 | sP175 | sP174 | sP173 | sP172 | sP171 | sP170 | sP169 | sP168 | sP167 | sP166 | sP165 | sP164 | sP163 | sP162 | sP161 | sP160 | sP159 | sP158 | sP157 | sP156 | sP155 | sP154 | sP153 | sP152 | sP151 | sP150 | sP149 | sP148 | sP147 | sP146 | sP145 | sP144 | sP143 | sP142 | sP141 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2789, f1459])).
fof(f1459, plain, ~ sP179, inference(subsumption_resolution, [], [f1107, f1084])).
fof(f1107, plain, (~ (e3 = op(unit, e3)) | ~ sP179), inference(definition_unfolding, [], [f481, f425])).
fof(f481, plain, (~ (e3 = op(e0, e3)) | ~ sP179), inference(cnf_transformation, [], [f210])).
fof(f210, plain, ((~ (e3 = op(e4, e3)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e0, e3))) | ~ sP179), inference(nnf_transformation, [], [f186])).
fof(f186, plain, ((~ (e3 = op(e4, e3)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e0, e3))) | ~ sP179), inference(usedef, [], [e186])).
fof(e186, plain, (sP179 <=> (~ (e3 = op(e4, e3)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP179])])).
fof(f2789, plain, (sP179 | sP178 | sP177 | sP176 | sP175 | sP174 | sP173 | sP172 | sP171 | sP170 | sP169 | sP168 | sP167 | sP166 | sP165 | sP164 | sP163 | sP162 | sP161 | sP160 | sP159 | sP158 | sP157 | sP156 | sP155 | sP154 | sP153 | sP152 | sP151 | sP150 | sP149 | sP148 | sP147 | sP146 | sP145 | sP144 | sP143 | sP142 | sP141 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2788, f1454])).
fof(f1454, plain, ~ sP180, inference(subsumption_resolution, [], [f479, f418])).
fof(f479, plain, (~ (e4 = op(e3, e3)) | ~ sP180), inference(cnf_transformation, [], [f209])).
fof(f209, plain, ((~ (e4 = op(e3, e4)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e3, e0))) | ~ sP180), inference(nnf_transformation, [], [f187])).
fof(f187, plain, ((~ (e4 = op(e3, e4)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e3, e0))) | ~ sP180), inference(usedef, [], [e187])).
fof(e187, plain, (sP180 <=> (~ (e4 = op(e3, e4)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP180])])).
fof(f2788, plain, (sP180 | sP179 | sP178 | sP177 | sP176 | sP175 | sP174 | sP173 | sP172 | sP171 | sP170 | sP169 | sP168 | sP167 | sP166 | sP165 | sP164 | sP163 | sP162 | sP161 | sP160 | sP159 | sP158 | sP157 | sP156 | sP155 | sP154 | sP153 | sP152 | sP151 | sP150 | sP149 | sP148 | sP147 | sP146 | sP145 | sP144 | sP143 | sP142 | sP141 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2787, f1452])).
fof(f1452, plain, ~ sP181, inference(subsumption_resolution, [], [f474, f418])).
fof(f474, plain, (~ (e4 = op(e3, e3)) | ~ sP181), inference(cnf_transformation, [], [f208])).
fof(f208, plain, ((~ (e4 = op(e4, e3)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e0, e3))) | ~ sP181), inference(nnf_transformation, [], [f188])).
fof(f188, plain, ((~ (e4 = op(e4, e3)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e0, e3))) | ~ sP181), inference(usedef, [], [e188])).
fof(e188, plain, (sP181 <=> (~ (e4 = op(e4, e3)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP181])])).
fof(f2787, plain, (sP181 | sP180 | sP179 | sP178 | sP177 | sP176 | sP175 | sP174 | sP173 | sP172 | sP171 | sP170 | sP169 | sP168 | sP167 | sP166 | sP165 | sP164 | sP163 | sP162 | sP161 | sP160 | sP159 | sP158 | sP157 | sP156 | sP155 | sP154 | sP153 | sP152 | sP151 | sP150 | sP149 | sP148 | sP147 | sP146 | sP145 | sP144 | sP143 | sP142 | sP141 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2786, f1450])).
fof(f1450, plain, ~ sP182, inference(subsumption_resolution, [], [f1102, f1075])).
fof(f1102, plain, (~ (op(e4, e2) = unit) | ~ sP182), inference(definition_unfolding, [], [f468, f425])).
fof(f468, plain, (~ (e0 = op(e4, e2)) | ~ sP182), inference(cnf_transformation, [], [f207])).
fof(f207, plain, ((~ (e0 = op(e4, e4)) & ~ (e0 = op(e4, e3)) & ~ (e0 = op(e4, e2)) & ~ (e0 = op(e4, e1)) & ~ (e0 = op(e4, e0))) | ~ sP182), inference(nnf_transformation, [], [f189])).
fof(f189, plain, ((~ (e0 = op(e4, e4)) & ~ (e0 = op(e4, e3)) & ~ (e0 = op(e4, e2)) & ~ (e0 = op(e4, e1)) & ~ (e0 = op(e4, e0))) | ~ sP182), inference(usedef, [], [e189])).
fof(e189, plain, (sP182 <=> (~ (e0 = op(e4, e4)) & ~ (e0 = op(e4, e3)) & ~ (e0 = op(e4, e2)) & ~ (e0 = op(e4, e1)) & ~ (e0 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP182])])).
fof(f2786, plain, (sP182 | sP181 | sP180 | sP179 | sP178 | sP177 | sP176 | sP175 | sP174 | sP173 | sP172 | sP171 | sP170 | sP169 | sP168 | sP167 | sP166 | sP165 | sP164 | sP163 | sP162 | sP161 | sP160 | sP159 | sP158 | sP157 | sP156 | sP155 | sP154 | sP153 | sP152 | sP151 | sP150 | sP149 | sP148 | sP147 | sP146 | sP145 | sP144 | sP143 | sP142 | sP141 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2785, f1447])).
fof(f1447, plain, ~ sP183, inference(subsumption_resolution, [], [f1096, f1077])).
fof(f1096, plain, (~ (op(e3, e4) = unit) | ~ sP183), inference(definition_unfolding, [], [f464, f425])).
fof(f464, plain, (~ (e0 = op(e3, e4)) | ~ sP183), inference(cnf_transformation, [], [f206])).
fof(f206, plain, ((~ (e0 = op(e4, e4)) & ~ (e0 = op(e3, e4)) & ~ (e0 = op(e2, e4)) & ~ (e0 = op(e1, e4)) & ~ (e0 = op(e0, e4))) | ~ sP183), inference(nnf_transformation, [], [f190])).
fof(f190, plain, ((~ (e0 = op(e4, e4)) & ~ (e0 = op(e3, e4)) & ~ (e0 = op(e2, e4)) & ~ (e0 = op(e1, e4)) & ~ (e0 = op(e0, e4))) | ~ sP183), inference(usedef, [], [e190])).
fof(e190, plain, (sP183 <=> (~ (e0 = op(e4, e4)) & ~ (e0 = op(e3, e4)) & ~ (e0 = op(e2, e4)) & ~ (e0 = op(e1, e4)) & ~ (e0 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP183])])).
fof(f2785, plain, (sP183 | sP182 | sP181 | sP180 | sP179 | sP178 | sP177 | sP176 | sP175 | sP174 | sP173 | sP172 | sP171 | sP170 | sP169 | sP168 | sP167 | sP166 | sP165 | sP164 | sP163 | sP162 | sP161 | sP160 | sP159 | sP158 | sP157 | sP156 | sP155 | sP154 | sP153 | sP152 | sP151 | sP150 | sP149 | sP148 | sP147 | sP146 | sP145 | sP144 | sP143 | sP142 | sP141 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2784, f1445])).
fof(f1445, plain, ~ sP184, inference(subsumption_resolution, [], [f459, f423])).
fof(f459, plain, (~ (e1 = op(e4, e3)) | ~ sP184), inference(cnf_transformation, [], [f205])).
fof(f205, plain, ((~ (e1 = op(e4, e4)) & ~ (e1 = op(e4, e3)) & ~ (e1 = op(e4, e2)) & ~ (e1 = op(e4, e1)) & ~ (e1 = op(e4, e0))) | ~ sP184), inference(nnf_transformation, [], [f191])).
fof(f191, plain, ((~ (e1 = op(e4, e4)) & ~ (e1 = op(e4, e3)) & ~ (e1 = op(e4, e2)) & ~ (e1 = op(e4, e1)) & ~ (e1 = op(e4, e0))) | ~ sP184), inference(usedef, [], [e191])).
fof(e191, plain, (sP184 <=> (~ (e1 = op(e4, e4)) & ~ (e1 = op(e4, e3)) & ~ (e1 = op(e4, e2)) & ~ (e1 = op(e4, e1)) & ~ (e1 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP184])])).
fof(f2784, plain, (sP184 | sP183 | sP182 | sP181 | sP180 | sP179 | sP178 | sP177 | sP176 | sP175 | sP174 | sP173 | sP172 | sP171 | sP170 | sP169 | sP168 | sP167 | sP166 | sP165 | sP164 | sP163 | sP162 | sP161 | sP160 | sP159 | sP158 | sP157 | sP156 | sP155 | sP154 | sP153 | sP152 | sP151 | sP150 | sP149 | sP148 | sP147 | sP146 | sP145 | sP144 | sP143 | sP142 | sP141 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2783, f1443])).
fof(f1443, plain, ~ sP185, inference(subsumption_resolution, [], [f453, f414])).
fof(f453, plain, (~ (e1 = op(e2, e4)) | ~ sP185), inference(cnf_transformation, [], [f204])).
fof(f204, plain, ((~ (e1 = op(e4, e4)) & ~ (e1 = op(e3, e4)) & ~ (e1 = op(e2, e4)) & ~ (e1 = op(e1, e4)) & ~ (e1 = op(e0, e4))) | ~ sP185), inference(nnf_transformation, [], [f192])).
fof(f192, plain, ((~ (e1 = op(e4, e4)) & ~ (e1 = op(e3, e4)) & ~ (e1 = op(e2, e4)) & ~ (e1 = op(e1, e4)) & ~ (e1 = op(e0, e4))) | ~ sP185), inference(usedef, [], [e192])).
fof(e192, plain, (sP185 <=> (~ (e1 = op(e4, e4)) & ~ (e1 = op(e3, e4)) & ~ (e1 = op(e2, e4)) & ~ (e1 = op(e1, e4)) & ~ (e1 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP185])])).
fof(f2783, plain, (sP185 | sP184 | sP183 | sP182 | sP181 | sP180 | sP179 | sP178 | sP177 | sP176 | sP175 | sP174 | sP173 | sP172 | sP171 | sP170 | sP169 | sP168 | sP167 | sP166 | sP165 | sP164 | sP163 | sP162 | sP161 | sP160 | sP159 | sP158 | sP157 | sP156 | sP155 | sP154 | sP153 | sP152 | sP151 | sP150 | sP149 | sP148 | sP147 | sP146 | sP145 | sP144 | sP143 | sP142 | sP141 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2782, f1440])).
fof(f1440, plain, ~ sP186, inference(subsumption_resolution, [], [f450, f424])).
fof(f450, plain, (~ (e2 = op(e4, e4)) | ~ sP186), inference(cnf_transformation, [], [f203])).
fof(f203, plain, ((~ (e2 = op(e4, e4)) & ~ (e2 = op(e4, e3)) & ~ (e2 = op(e4, e2)) & ~ (e2 = op(e4, e1)) & ~ (e2 = op(e4, e0))) | ~ sP186), inference(nnf_transformation, [], [f193])).
fof(f193, plain, ((~ (e2 = op(e4, e4)) & ~ (e2 = op(e4, e3)) & ~ (e2 = op(e4, e2)) & ~ (e2 = op(e4, e1)) & ~ (e2 = op(e4, e0))) | ~ sP186), inference(usedef, [], [e193])).
fof(e193, plain, (sP186 <=> (~ (e2 = op(e4, e4)) & ~ (e2 = op(e4, e3)) & ~ (e2 = op(e4, e2)) & ~ (e2 = op(e4, e1)) & ~ (e2 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP186])])).
fof(f2782, plain, (sP186 | sP185 | sP184 | sP183 | sP182 | sP181 | sP180 | sP179 | sP178 | sP177 | sP176 | sP175 | sP174 | sP173 | sP172 | sP171 | sP170 | sP169 | sP168 | sP167 | sP166 | sP165 | sP164 | sP163 | sP162 | sP161 | sP160 | sP159 | sP158 | sP157 | sP156 | sP155 | sP154 | sP153 | sP152 | sP151 | sP150 | sP149 | sP148 | sP147 | sP146 | sP145 | sP144 | sP143 | sP142 | sP141 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2781, f1439])).
fof(f1439, plain, ~ sP187, inference(subsumption_resolution, [], [f445, f424])).
fof(f445, plain, (~ (e2 = op(e4, e4)) | ~ sP187), inference(cnf_transformation, [], [f202])).
fof(f202, plain, ((~ (e2 = op(e4, e4)) & ~ (e2 = op(e3, e4)) & ~ (e2 = op(e2, e4)) & ~ (e2 = op(e1, e4)) & ~ (e2 = op(e0, e4))) | ~ sP187), inference(nnf_transformation, [], [f194])).
fof(f194, plain, ((~ (e2 = op(e4, e4)) & ~ (e2 = op(e3, e4)) & ~ (e2 = op(e2, e4)) & ~ (e2 = op(e1, e4)) & ~ (e2 = op(e0, e4))) | ~ sP187), inference(usedef, [], [e194])).
fof(e194, plain, (sP187 <=> (~ (e2 = op(e4, e4)) & ~ (e2 = op(e3, e4)) & ~ (e2 = op(e2, e4)) & ~ (e2 = op(e1, e4)) & ~ (e2 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP187])])).
fof(f2781, plain, (sP187 | sP186 | sP185 | sP184 | sP183 | sP182 | sP181 | sP180 | sP179 | sP178 | sP177 | sP176 | sP175 | sP174 | sP173 | sP172 | sP171 | sP170 | sP169 | sP168 | sP167 | sP166 | sP165 | sP164 | sP163 | sP162 | sP161 | sP160 | sP159 | sP158 | sP157 | sP156 | sP155 | sP154 | sP153 | sP152 | sP151 | sP150 | sP149 | sP148 | sP147 | sP146 | sP145 | sP144 | sP143 | sP142 | sP141 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2780, f1438])).
fof(f1438, plain, ~ sP188, inference(subsumption_resolution, [], [f437, f421])).
fof(f437, plain, (~ (e3 = op(e4, e1)) | ~ sP188), inference(cnf_transformation, [], [f201])).
fof(f201, plain, ((~ (e3 = op(e4, e4)) & ~ (e3 = op(e4, e3)) & ~ (e3 = op(e4, e2)) & ~ (e3 = op(e4, e1)) & ~ (e3 = op(e4, e0))) | ~ sP188), inference(nnf_transformation, [], [f195])).
fof(f195, plain, ((~ (e3 = op(e4, e4)) & ~ (e3 = op(e4, e3)) & ~ (e3 = op(e4, e2)) & ~ (e3 = op(e4, e1)) & ~ (e3 = op(e4, e0))) | ~ sP188), inference(usedef, [], [e195])).
fof(e195, plain, (sP188 <=> (~ (e3 = op(e4, e4)) & ~ (e3 = op(e4, e3)) & ~ (e3 = op(e4, e2)) & ~ (e3 = op(e4, e1)) & ~ (e3 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP188])])).
fof(f2780, plain, (sP188 | sP187 | sP186 | sP185 | sP184 | sP183 | sP182 | sP181 | sP180 | sP179 | sP178 | sP177 | sP176 | sP175 | sP174 | sP173 | sP172 | sP171 | sP170 | sP169 | sP168 | sP167 | sP166 | sP165 | sP164 | sP163 | sP162 | sP161 | sP160 | sP159 | sP158 | sP157 | sP156 | sP155 | sP154 | sP153 | sP152 | sP151 | sP150 | sP149 | sP148 | sP147 | sP146 | sP145 | sP144 | sP143 | sP142 | sP141 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2779, f1434])).
fof(f1434, plain, ~ sP189, inference(subsumption_resolution, [], [f432, f409])).
fof(f432, plain, (~ (e3 = op(e1, e4)) | ~ sP189), inference(cnf_transformation, [], [f200])).
fof(f200, plain, ((~ (e3 = op(e4, e4)) & ~ (e3 = op(e3, e4)) & ~ (e3 = op(e2, e4)) & ~ (e3 = op(e1, e4)) & ~ (e3 = op(e0, e4))) | ~ sP189), inference(nnf_transformation, [], [f196])).
fof(f196, plain, ((~ (e3 = op(e4, e4)) & ~ (e3 = op(e3, e4)) & ~ (e3 = op(e2, e4)) & ~ (e3 = op(e1, e4)) & ~ (e3 = op(e0, e4))) | ~ sP189), inference(usedef, [], [e196])).
fof(e196, plain, (sP189 <=> (~ (e3 = op(e4, e4)) & ~ (e3 = op(e3, e4)) & ~ (e3 = op(e2, e4)) & ~ (e3 = op(e1, e4)) & ~ (e3 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP189])])).
fof(f2779, plain, (sP189 | sP188 | sP187 | sP186 | sP185 | sP184 | sP183 | sP182 | sP181 | sP180 | sP179 | sP178 | sP177 | sP176 | sP175 | sP174 | sP173 | sP172 | sP171 | sP170 | sP169 | sP168 | sP167 | sP166 | sP165 | sP164 | sP163 | sP162 | sP161 | sP160 | sP159 | sP158 | sP157 | sP156 | sP155 | sP154 | sP153 | sP152 | sP151 | sP150 | sP149 | sP148 | sP147 | sP146 | sP145 | sP144 | sP143 | sP142 | sP141 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f2778, f1430])).
fof(f1430, plain, ~ sP190, inference(subsumption_resolution, [], [f1088, f1076])).
fof(f1088, plain, (~ (e4 = op(e4, unit)) | ~ sP190), inference(definition_unfolding, [], [f426, f425])).
fof(f426, plain, (~ (e4 = op(e4, e0)) | ~ sP190), inference(cnf_transformation, [], [f199])).
fof(f199, plain, ((~ (e4 = op(e4, e4)) & ~ (e4 = op(e4, e3)) & ~ (e4 = op(e4, e2)) & ~ (e4 = op(e4, e1)) & ~ (e4 = op(e4, e0))) | ~ sP190), inference(nnf_transformation, [], [f197])).
fof(f197, plain, ((~ (e4 = op(e4, e4)) & ~ (e4 = op(e4, e3)) & ~ (e4 = op(e4, e2)) & ~ (e4 = op(e4, e1)) & ~ (e4 = op(e4, e0))) | ~ sP190), inference(usedef, [], [e197])).
fof(e197, plain, (sP190 <=> (~ (e4 = op(e4, e4)) & ~ (e4 = op(e4, e3)) & ~ (e4 = op(e4, e2)) & ~ (e4 = op(e4, e1)) & ~ (e4 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP190])])).
fof(f2778, plain, (sP190 | sP189 | sP188 | sP187 | sP186 | sP185 | sP184 | sP183 | sP182 | sP181 | sP180 | sP179 | sP178 | sP177 | sP176 | sP175 | sP174 | sP173 | sP172 | sP171 | sP170 | sP169 | sP168 | sP167 | sP166 | sP165 | sP164 | sP163 | sP162 | sP161 | sP160 | sP159 | sP158 | sP157 | sP156 | sP155 | sP154 | sP153 | sP152 | sP151 | sP150 | sP149 | sP148 | sP147 | sP146 | sP145 | sP144 | sP143 | sP142 | sP141 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(subsumption_resolution, [], [f1400, f1083])).
fof(f1400, plain, (~ (e4 = op(unit, e4)) | sP190 | sP189 | sP188 | sP187 | sP186 | sP185 | sP184 | sP183 | sP182 | sP181 | sP180 | sP179 | sP178 | sP177 | sP176 | sP175 | sP174 | sP173 | sP172 | sP171 | sP170 | sP169 | sP168 | sP167 | sP166 | sP165 | sP164 | sP163 | sP162 | sP161 | sP160 | sP159 | sP158 | sP157 | sP156 | sP155 | sP154 | sP153 | sP152 | sP151 | sP150 | sP149 | sP148 | sP147 | sP146 | sP145 | sP144 | sP143 | sP142 | sP141 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(duplicate_literal_removal, [], [f1399])).
fof(f1399, plain, (~ (e4 = op(unit, e4)) | sP190 | sP189 | sP188 | sP187 | sP186 | sP185 | sP184 | sP183 | sP182 | sP181 | sP180 | sP179 | sP178 | sP177 | sP176 | sP175 | sP174 | sP173 | sP172 | sP171 | sP170 | sP169 | sP168 | sP167 | sP166 | sP165 | sP164 | sP163 | sP162 | sP161 | sP160 | sP159 | sP158 | sP157 | sP156 | sP155 | sP154 | sP153 | sP152 | sP151 | sP150 | sP149 | sP148 | sP147 | sP146 | sP145 | sP144 | sP143 | sP142 | sP141 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | ~ (unit = op(unit, unit)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(definition_unfolding, [], [f1066, f425, f425, f425, f425, f425])).
fof(f1066, plain, (~ (e4 = op(e0, e4)) | sP190 | sP189 | sP188 | sP187 | sP186 | sP185 | sP184 | sP183 | sP182 | sP181 | sP180 | sP179 | sP178 | sP177 | sP176 | sP175 | sP174 | sP173 | sP172 | sP171 | sP170 | sP169 | sP168 | sP167 | sP166 | sP165 | sP164 | sP163 | sP162 | sP161 | sP160 | sP159 | sP158 | sP157 | sP156 | sP155 | sP154 | sP153 | sP152 | sP151 | sP150 | sP149 | sP148 | sP147 | sP146 | sP145 | sP144 | sP143 | sP142 | sP141 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (e0 = op(e0, unit)) | ~ (e0 = op(unit, e0)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(cnf_transformation, [], [f198])).
fof(f198, plain, ((~ (e4 = op(e4, e4)) & ~ (e4 = op(e3, e4)) & ~ (e4 = op(e2, e4)) & ~ (e4 = op(e1, e4)) & ~ (e4 = op(e0, e4))) | sP190 | sP189 | sP188 | sP187 | sP186 | sP185 | sP184 | sP183 | sP182 | sP181 | sP180 | sP179 | sP178 | sP177 | sP176 | sP175 | sP174 | sP173 | sP172 | sP171 | sP170 | sP169 | sP168 | sP167 | sP166 | sP165 | sP164 | sP163 | sP162 | sP161 | sP160 | sP159 | sP158 | sP157 | sP156 | sP155 | sP154 | sP153 | sP152 | sP151 | sP150 | sP149 | sP148 | sP147 | sP146 | sP145 | sP144 | sP143 | sP142 | sP141 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (e0 = op(e0, unit)) | ~ (e0 = op(unit, e0)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115), inference(definition_folding, [], [f6, e197, e196, e195, e194, e193, e192, e191, e190, e189, e188, e187, e186, e185, e184, e183, e182, e181, e180, e179, e178, e177, e176, e175, e174, e173, e172, e171, e170, e169, e168, e167, e166, e165, e164, e163, e162, e161, e160, e159, e158, e157, e156, e155, e154, e153, e152, e151, e150, e149, e148, e147, e146, e145, e144, e143, e142, e141, e140, e139, e138, e137, e136, e135, e134, e133, e132, e131, e130, e129, e128, e127, e126, e125, e124, e123, e122, e121, e120, e119, e118, e117, e116, e115, e114, e113, e112, e111, e110, e109, e108, e107, e106, e105, e104, e103, e102, e101, e100, e99, e98, e97, e96, e95, e94, e93, e92, e91, e90, e89, e88, e87, e86, e85, e84, e83, e82, e81, e80, e79, e78, e77, e76, e75, e74, e73, e72, e71, e70, e69, e68, e67, e66, e65, e64, e63, e62, e61, e60, e59, e58, e57, e56, e55, e54, e53, e52, e51, e50, e49, e48, e47, e46, e45, e44, e43, e42, e41, e40, e39, e38, e37, e36, e35, e34, e33, e32, e31, e30, e29, e28, e27, e26, e25, e24, e23, e22, e21, e20, e19, e18, e17, e16, e15, e14, e13, e12, e11, e10, e9, e8, e7])).
fof(f7, plain, ((((e0 = unit) | ~ (e0 = op(e0, e0))) & (e0 = op(e0, e0))) | ~ sP0), inference(usedef, [], [e7])).
fof(e7, plain, (sP0 <=> (((e0 = unit) | ~ (e0 = op(e0, e0))) & (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f8, plain, ((((e0 = unit) | ~ (e1 = op(e0, e0))) & (e0 = op(e0, e1))) | ~ sP1), inference(usedef, [], [e8])).
fof(e8, plain, (sP1 <=> (((e0 = unit) | ~ (e1 = op(e0, e0))) & (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f9, plain, ((((e0 = unit) | ~ (e2 = op(e0, e0))) & (e0 = op(e0, e2))) | ~ sP2), inference(usedef, [], [e9])).
fof(e9, plain, (sP2 <=> (((e0 = unit) | ~ (e2 = op(e0, e0))) & (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f10, plain, ((((e0 = unit) | ~ (e3 = op(e0, e0))) & (e0 = op(e0, e3))) | ~ sP3), inference(usedef, [], [e10])).
fof(e10, plain, (sP3 <=> (((e0 = unit) | ~ (e3 = op(e0, e0))) & (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f11, plain, ((((e0 = unit) | ~ (e4 = op(e0, e0))) & (e0 = op(e0, e4))) | ~ sP4), inference(usedef, [], [e11])).
fof(e11, plain, (sP4 <=> (((e0 = unit) | ~ (e4 = op(e0, e0))) & (e0 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f12, plain, ((((e0 = unit) | ~ (e0 = op(e1, e0))) & (e0 = op(e1, e0))) | ~ sP5), inference(usedef, [], [e12])).
fof(e12, plain, (sP5 <=> (((e0 = unit) | ~ (e0 = op(e1, e0))) & (e0 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f13, plain, ((((e0 = unit) | ~ (e1 = op(e1, e0))) & (e0 = op(e1, e1))) | ~ sP6), inference(usedef, [], [e13])).
fof(e13, plain, (sP6 <=> (((e0 = unit) | ~ (e1 = op(e1, e0))) & (e0 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f14, plain, ((((e0 = unit) | ~ (e2 = op(e1, e0))) & (e0 = op(e1, e2))) | ~ sP7), inference(usedef, [], [e14])).
fof(e14, plain, (sP7 <=> (((e0 = unit) | ~ (e2 = op(e1, e0))) & (e0 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f15, plain, ((((e0 = unit) | ~ (e3 = op(e1, e0))) & (e0 = op(e1, e3))) | ~ sP8), inference(usedef, [], [e15])).
fof(e15, plain, (sP8 <=> (((e0 = unit) | ~ (e3 = op(e1, e0))) & (e0 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f16, plain, ((((e0 = unit) | ~ (e4 = op(e1, e0))) & (e0 = op(e1, e4))) | ~ sP9), inference(usedef, [], [e16])).
fof(e16, plain, (sP9 <=> (((e0 = unit) | ~ (e4 = op(e1, e0))) & (e0 = op(e1, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f17, plain, ((((e0 = unit) | ~ (e0 = op(e2, e0))) & (e0 = op(e2, e0))) | ~ sP10), inference(usedef, [], [e17])).
fof(e17, plain, (sP10 <=> (((e0 = unit) | ~ (e0 = op(e2, e0))) & (e0 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f18, plain, ((((e0 = unit) | ~ (e1 = op(e2, e0))) & (e0 = op(e2, e1))) | ~ sP11), inference(usedef, [], [e18])).
fof(e18, plain, (sP11 <=> (((e0 = unit) | ~ (e1 = op(e2, e0))) & (e0 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f19, plain, ((((e0 = unit) | ~ (e2 = op(e2, e0))) & (e0 = op(e2, e2))) | ~ sP12), inference(usedef, [], [e19])).
fof(e19, plain, (sP12 <=> (((e0 = unit) | ~ (e2 = op(e2, e0))) & (e0 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f20, plain, ((((e0 = unit) | ~ (e3 = op(e2, e0))) & (e0 = op(e2, e3))) | ~ sP13), inference(usedef, [], [e20])).
fof(e20, plain, (sP13 <=> (((e0 = unit) | ~ (e3 = op(e2, e0))) & (e0 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f21, plain, ((((e0 = unit) | ~ (e4 = op(e2, e0))) & (e0 = op(e2, e4))) | ~ sP14), inference(usedef, [], [e21])).
fof(e21, plain, (sP14 <=> (((e0 = unit) | ~ (e4 = op(e2, e0))) & (e0 = op(e2, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f22, plain, ((((e0 = unit) | ~ (e0 = op(e3, e0))) & (e0 = op(e3, e0))) | ~ sP15), inference(usedef, [], [e22])).
fof(e22, plain, (sP15 <=> (((e0 = unit) | ~ (e0 = op(e3, e0))) & (e0 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f23, plain, ((((e0 = unit) | ~ (e1 = op(e3, e0))) & (e0 = op(e3, e1))) | ~ sP16), inference(usedef, [], [e23])).
fof(e23, plain, (sP16 <=> (((e0 = unit) | ~ (e1 = op(e3, e0))) & (e0 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f24, plain, ((((e0 = unit) | ~ (e2 = op(e3, e0))) & (e0 = op(e3, e2))) | ~ sP17), inference(usedef, [], [e24])).
fof(e24, plain, (sP17 <=> (((e0 = unit) | ~ (e2 = op(e3, e0))) & (e0 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f25, plain, ((((e0 = unit) | ~ (e3 = op(e3, e0))) & (e0 = op(e3, e3))) | ~ sP18), inference(usedef, [], [e25])).
fof(e25, plain, (sP18 <=> (((e0 = unit) | ~ (e3 = op(e3, e0))) & (e0 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f26, plain, ((((e0 = unit) | ~ (e4 = op(e3, e0))) & (e0 = op(e3, e4))) | ~ sP19), inference(usedef, [], [e26])).
fof(e26, plain, (sP19 <=> (((e0 = unit) | ~ (e4 = op(e3, e0))) & (e0 = op(e3, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP19])])).
fof(f27, plain, ((((e0 = unit) | ~ (e0 = op(e4, e0))) & (e0 = op(e4, e0))) | ~ sP20), inference(usedef, [], [e27])).
fof(e27, plain, (sP20 <=> (((e0 = unit) | ~ (e0 = op(e4, e0))) & (e0 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP20])])).
fof(f28, plain, ((((e0 = unit) | ~ (e1 = op(e4, e0))) & (e0 = op(e4, e1))) | ~ sP21), inference(usedef, [], [e28])).
fof(e28, plain, (sP21 <=> (((e0 = unit) | ~ (e1 = op(e4, e0))) & (e0 = op(e4, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP21])])).
fof(f29, plain, ((((e1 = unit) | ~ (e0 = op(e0, e1))) & (e1 = op(e0, e0))) | ~ sP22), inference(usedef, [], [e29])).
fof(e29, plain, (sP22 <=> (((e1 = unit) | ~ (e0 = op(e0, e1))) & (e1 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP22])])).
fof(f30, plain, ((((e1 = unit) | ~ (e1 = op(e0, e1))) & (e1 = op(e0, e1))) | ~ sP23), inference(usedef, [], [e30])).
fof(e30, plain, (sP23 <=> (((e1 = unit) | ~ (e1 = op(e0, e1))) & (e1 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP23])])).
fof(f31, plain, ((((e1 = unit) | ~ (e2 = op(e0, e1))) & (e1 = op(e0, e2))) | ~ sP24), inference(usedef, [], [e31])).
fof(e31, plain, (sP24 <=> (((e1 = unit) | ~ (e2 = op(e0, e1))) & (e1 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP24])])).
fof(f32, plain, ((((e1 = unit) | ~ (e3 = op(e0, e1))) & (e1 = op(e0, e3))) | ~ sP25), inference(usedef, [], [e32])).
fof(e32, plain, (sP25 <=> (((e1 = unit) | ~ (e3 = op(e0, e1))) & (e1 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP25])])).
fof(f33, plain, ((((e1 = unit) | ~ (e4 = op(e0, e1))) & (e1 = op(e0, e4))) | ~ sP26), inference(usedef, [], [e33])).
fof(e33, plain, (sP26 <=> (((e1 = unit) | ~ (e4 = op(e0, e1))) & (e1 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP26])])).
fof(f34, plain, ((((e1 = unit) | ~ (e0 = op(e1, e1))) & (e1 = op(e1, e0))) | ~ sP27), inference(usedef, [], [e34])).
fof(e34, plain, (sP27 <=> (((e1 = unit) | ~ (e0 = op(e1, e1))) & (e1 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP27])])).
fof(f35, plain, ((((e1 = unit) | ~ (e1 = op(e1, e1))) & (e1 = op(e1, e1))) | ~ sP28), inference(usedef, [], [e35])).
fof(e35, plain, (sP28 <=> (((e1 = unit) | ~ (e1 = op(e1, e1))) & (e1 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP28])])).
fof(f36, plain, ((((e1 = unit) | ~ (e2 = op(e1, e1))) & (e1 = op(e1, e2))) | ~ sP29), inference(usedef, [], [e36])).
fof(e36, plain, (sP29 <=> (((e1 = unit) | ~ (e2 = op(e1, e1))) & (e1 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP29])])).
fof(f37, plain, ((((e1 = unit) | ~ (e3 = op(e1, e1))) & (e1 = op(e1, e3))) | ~ sP30), inference(usedef, [], [e37])).
fof(e37, plain, (sP30 <=> (((e1 = unit) | ~ (e3 = op(e1, e1))) & (e1 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP30])])).
fof(f38, plain, ((((e1 = unit) | ~ (e4 = op(e1, e1))) & (e1 = op(e1, e4))) | ~ sP31), inference(usedef, [], [e38])).
fof(e38, plain, (sP31 <=> (((e1 = unit) | ~ (e4 = op(e1, e1))) & (e1 = op(e1, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP31])])).
fof(f39, plain, ((((e1 = unit) | ~ (e0 = op(e2, e1))) & (e1 = op(e2, e0))) | ~ sP32), inference(usedef, [], [e39])).
fof(e39, plain, (sP32 <=> (((e1 = unit) | ~ (e0 = op(e2, e1))) & (e1 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP32])])).
fof(f40, plain, ((((e1 = unit) | ~ (e1 = op(e2, e1))) & (e1 = op(e2, e1))) | ~ sP33), inference(usedef, [], [e40])).
fof(e40, plain, (sP33 <=> (((e1 = unit) | ~ (e1 = op(e2, e1))) & (e1 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP33])])).
fof(f41, plain, ((((e1 = unit) | ~ (e2 = op(e2, e1))) & (e1 = op(e2, e2))) | ~ sP34), inference(usedef, [], [e41])).
fof(e41, plain, (sP34 <=> (((e1 = unit) | ~ (e2 = op(e2, e1))) & (e1 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP34])])).
fof(f42, plain, ((((e1 = unit) | ~ (e3 = op(e2, e1))) & (e1 = op(e2, e3))) | ~ sP35), inference(usedef, [], [e42])).
fof(e42, plain, (sP35 <=> (((e1 = unit) | ~ (e3 = op(e2, e1))) & (e1 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP35])])).
fof(f43, plain, ((((e1 = unit) | ~ (e4 = op(e2, e1))) & (e1 = op(e2, e4))) | ~ sP36), inference(usedef, [], [e43])).
fof(e43, plain, (sP36 <=> (((e1 = unit) | ~ (e4 = op(e2, e1))) & (e1 = op(e2, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP36])])).
fof(f44, plain, ((((e1 = unit) | ~ (e0 = op(e3, e1))) & (e1 = op(e3, e0))) | ~ sP37), inference(usedef, [], [e44])).
fof(e44, plain, (sP37 <=> (((e1 = unit) | ~ (e0 = op(e3, e1))) & (e1 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP37])])).
fof(f45, plain, ((((e1 = unit) | ~ (e1 = op(e3, e1))) & (e1 = op(e3, e1))) | ~ sP38), inference(usedef, [], [e45])).
fof(e45, plain, (sP38 <=> (((e1 = unit) | ~ (e1 = op(e3, e1))) & (e1 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP38])])).
fof(f46, plain, ((((e1 = unit) | ~ (e2 = op(e3, e1))) & (e1 = op(e3, e2))) | ~ sP39), inference(usedef, [], [e46])).
fof(e46, plain, (sP39 <=> (((e1 = unit) | ~ (e2 = op(e3, e1))) & (e1 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP39])])).
fof(f47, plain, ((((e1 = unit) | ~ (e3 = op(e3, e1))) & (e1 = op(e3, e3))) | ~ sP40), inference(usedef, [], [e47])).
fof(e47, plain, (sP40 <=> (((e1 = unit) | ~ (e3 = op(e3, e1))) & (e1 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP40])])).
fof(f48, plain, ((((e1 = unit) | ~ (e4 = op(e3, e1))) & (e1 = op(e3, e4))) | ~ sP41), inference(usedef, [], [e48])).
fof(e48, plain, (sP41 <=> (((e1 = unit) | ~ (e4 = op(e3, e1))) & (e1 = op(e3, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP41])])).
fof(f49, plain, ((((e1 = unit) | ~ (e0 = op(e4, e1))) & (e1 = op(e4, e0))) | ~ sP42), inference(usedef, [], [e49])).
fof(e49, plain, (sP42 <=> (((e1 = unit) | ~ (e0 = op(e4, e1))) & (e1 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP42])])).
fof(f50, plain, ((((e1 = unit) | ~ (e1 = op(e4, e1))) & (e1 = op(e4, e1))) | ~ sP43), inference(usedef, [], [e50])).
fof(e50, plain, (sP43 <=> (((e1 = unit) | ~ (e1 = op(e4, e1))) & (e1 = op(e4, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP43])])).
fof(f51, plain, ((((e2 = unit) | ~ (e0 = op(e0, e2))) & (e2 = op(e0, e0))) | ~ sP44), inference(usedef, [], [e51])).
fof(e51, plain, (sP44 <=> (((e2 = unit) | ~ (e0 = op(e0, e2))) & (e2 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP44])])).
fof(f52, plain, ((((e2 = unit) | ~ (e1 = op(e0, e2))) & (e2 = op(e0, e1))) | ~ sP45), inference(usedef, [], [e52])).
fof(e52, plain, (sP45 <=> (((e2 = unit) | ~ (e1 = op(e0, e2))) & (e2 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP45])])).
fof(f53, plain, ((((e2 = unit) | ~ (e2 = op(e0, e2))) & (e2 = op(e0, e2))) | ~ sP46), inference(usedef, [], [e53])).
fof(e53, plain, (sP46 <=> (((e2 = unit) | ~ (e2 = op(e0, e2))) & (e2 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP46])])).
fof(f54, plain, ((((e2 = unit) | ~ (e3 = op(e0, e2))) & (e2 = op(e0, e3))) | ~ sP47), inference(usedef, [], [e54])).
fof(e54, plain, (sP47 <=> (((e2 = unit) | ~ (e3 = op(e0, e2))) & (e2 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP47])])).
fof(f55, plain, ((((e2 = unit) | ~ (e4 = op(e0, e2))) & (e2 = op(e0, e4))) | ~ sP48), inference(usedef, [], [e55])).
fof(e55, plain, (sP48 <=> (((e2 = unit) | ~ (e4 = op(e0, e2))) & (e2 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP48])])).
fof(f56, plain, ((((e2 = unit) | ~ (e0 = op(e1, e2))) & (e2 = op(e1, e0))) | ~ sP49), inference(usedef, [], [e56])).
fof(e56, plain, (sP49 <=> (((e2 = unit) | ~ (e0 = op(e1, e2))) & (e2 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP49])])).
fof(f57, plain, ((((e2 = unit) | ~ (e1 = op(e1, e2))) & (e2 = op(e1, e1))) | ~ sP50), inference(usedef, [], [e57])).
fof(e57, plain, (sP50 <=> (((e2 = unit) | ~ (e1 = op(e1, e2))) & (e2 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP50])])).
fof(f58, plain, ((((e2 = unit) | ~ (e2 = op(e1, e2))) & (e2 = op(e1, e2))) | ~ sP51), inference(usedef, [], [e58])).
fof(e58, plain, (sP51 <=> (((e2 = unit) | ~ (e2 = op(e1, e2))) & (e2 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP51])])).
fof(f59, plain, ((((e2 = unit) | ~ (e3 = op(e1, e2))) & (e2 = op(e1, e3))) | ~ sP52), inference(usedef, [], [e59])).
fof(e59, plain, (sP52 <=> (((e2 = unit) | ~ (e3 = op(e1, e2))) & (e2 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP52])])).
fof(f60, plain, ((((e2 = unit) | ~ (e4 = op(e1, e2))) & (e2 = op(e1, e4))) | ~ sP53), inference(usedef, [], [e60])).
fof(e60, plain, (sP53 <=> (((e2 = unit) | ~ (e4 = op(e1, e2))) & (e2 = op(e1, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP53])])).
fof(f61, plain, ((((e2 = unit) | ~ (e0 = op(e2, e2))) & (e2 = op(e2, e0))) | ~ sP54), inference(usedef, [], [e61])).
fof(e61, plain, (sP54 <=> (((e2 = unit) | ~ (e0 = op(e2, e2))) & (e2 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP54])])).
fof(f62, plain, ((((e2 = unit) | ~ (e1 = op(e2, e2))) & (e2 = op(e2, e1))) | ~ sP55), inference(usedef, [], [e62])).
fof(e62, plain, (sP55 <=> (((e2 = unit) | ~ (e1 = op(e2, e2))) & (e2 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP55])])).
fof(f63, plain, ((((e2 = unit) | ~ (e2 = op(e2, e2))) & (e2 = op(e2, e2))) | ~ sP56), inference(usedef, [], [e63])).
fof(e63, plain, (sP56 <=> (((e2 = unit) | ~ (e2 = op(e2, e2))) & (e2 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP56])])).
fof(f64, plain, ((((e2 = unit) | ~ (e3 = op(e2, e2))) & (e2 = op(e2, e3))) | ~ sP57), inference(usedef, [], [e64])).
fof(e64, plain, (sP57 <=> (((e2 = unit) | ~ (e3 = op(e2, e2))) & (e2 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP57])])).
fof(f65, plain, ((((e2 = unit) | ~ (e4 = op(e2, e2))) & (e2 = op(e2, e4))) | ~ sP58), inference(usedef, [], [e65])).
fof(e65, plain, (sP58 <=> (((e2 = unit) | ~ (e4 = op(e2, e2))) & (e2 = op(e2, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP58])])).
fof(f66, plain, ((((e2 = unit) | ~ (e0 = op(e3, e2))) & (e2 = op(e3, e0))) | ~ sP59), inference(usedef, [], [e66])).
fof(e66, plain, (sP59 <=> (((e2 = unit) | ~ (e0 = op(e3, e2))) & (e2 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP59])])).
fof(f67, plain, ((((e2 = unit) | ~ (e1 = op(e3, e2))) & (e2 = op(e3, e1))) | ~ sP60), inference(usedef, [], [e67])).
fof(e67, plain, (sP60 <=> (((e2 = unit) | ~ (e1 = op(e3, e2))) & (e2 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP60])])).
fof(f68, plain, ((((e2 = unit) | ~ (e2 = op(e3, e2))) & (e2 = op(e3, e2))) | ~ sP61), inference(usedef, [], [e68])).
fof(e68, plain, (sP61 <=> (((e2 = unit) | ~ (e2 = op(e3, e2))) & (e2 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP61])])).
fof(f69, plain, ((((e2 = unit) | ~ (e3 = op(e3, e2))) & (e2 = op(e3, e3))) | ~ sP62), inference(usedef, [], [e69])).
fof(e69, plain, (sP62 <=> (((e2 = unit) | ~ (e3 = op(e3, e2))) & (e2 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP62])])).
fof(f70, plain, ((((e2 = unit) | ~ (e4 = op(e3, e2))) & (e2 = op(e3, e4))) | ~ sP63), inference(usedef, [], [e70])).
fof(e70, plain, (sP63 <=> (((e2 = unit) | ~ (e4 = op(e3, e2))) & (e2 = op(e3, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP63])])).
fof(f71, plain, ((((e2 = unit) | ~ (e0 = op(e4, e2))) & (e2 = op(e4, e0))) | ~ sP64), inference(usedef, [], [e71])).
fof(e71, plain, (sP64 <=> (((e2 = unit) | ~ (e0 = op(e4, e2))) & (e2 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP64])])).
fof(f72, plain, ((((e2 = unit) | ~ (e1 = op(e4, e2))) & (e2 = op(e4, e1))) | ~ sP65), inference(usedef, [], [e72])).
fof(e72, plain, (sP65 <=> (((e2 = unit) | ~ (e1 = op(e4, e2))) & (e2 = op(e4, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP65])])).
fof(f73, plain, ((((e3 = unit) | ~ (e0 = op(e0, e3))) & (e3 = op(e0, e0))) | ~ sP66), inference(usedef, [], [e73])).
fof(e73, plain, (sP66 <=> (((e3 = unit) | ~ (e0 = op(e0, e3))) & (e3 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP66])])).
fof(f74, plain, ((((e3 = unit) | ~ (e1 = op(e0, e3))) & (e3 = op(e0, e1))) | ~ sP67), inference(usedef, [], [e74])).
fof(e74, plain, (sP67 <=> (((e3 = unit) | ~ (e1 = op(e0, e3))) & (e3 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP67])])).
fof(f75, plain, ((((e3 = unit) | ~ (e2 = op(e0, e3))) & (e3 = op(e0, e2))) | ~ sP68), inference(usedef, [], [e75])).
fof(e75, plain, (sP68 <=> (((e3 = unit) | ~ (e2 = op(e0, e3))) & (e3 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP68])])).
fof(f76, plain, ((((e3 = unit) | ~ (e3 = op(e0, e3))) & (e3 = op(e0, e3))) | ~ sP69), inference(usedef, [], [e76])).
fof(e76, plain, (sP69 <=> (((e3 = unit) | ~ (e3 = op(e0, e3))) & (e3 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP69])])).
fof(f77, plain, ((((e3 = unit) | ~ (e4 = op(e0, e3))) & (e3 = op(e0, e4))) | ~ sP70), inference(usedef, [], [e77])).
fof(e77, plain, (sP70 <=> (((e3 = unit) | ~ (e4 = op(e0, e3))) & (e3 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP70])])).
fof(f78, plain, ((((e3 = unit) | ~ (e0 = op(e1, e3))) & (e3 = op(e1, e0))) | ~ sP71), inference(usedef, [], [e78])).
fof(e78, plain, (sP71 <=> (((e3 = unit) | ~ (e0 = op(e1, e3))) & (e3 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP71])])).
fof(f79, plain, ((((e3 = unit) | ~ (e1 = op(e1, e3))) & (e3 = op(e1, e1))) | ~ sP72), inference(usedef, [], [e79])).
fof(e79, plain, (sP72 <=> (((e3 = unit) | ~ (e1 = op(e1, e3))) & (e3 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP72])])).
fof(f80, plain, ((((e3 = unit) | ~ (e2 = op(e1, e3))) & (e3 = op(e1, e2))) | ~ sP73), inference(usedef, [], [e80])).
fof(e80, plain, (sP73 <=> (((e3 = unit) | ~ (e2 = op(e1, e3))) & (e3 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP73])])).
fof(f81, plain, ((((e3 = unit) | ~ (e3 = op(e1, e3))) & (e3 = op(e1, e3))) | ~ sP74), inference(usedef, [], [e81])).
fof(e81, plain, (sP74 <=> (((e3 = unit) | ~ (e3 = op(e1, e3))) & (e3 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP74])])).
fof(f82, plain, ((((e3 = unit) | ~ (e4 = op(e1, e3))) & (e3 = op(e1, e4))) | ~ sP75), inference(usedef, [], [e82])).
fof(e82, plain, (sP75 <=> (((e3 = unit) | ~ (e4 = op(e1, e3))) & (e3 = op(e1, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP75])])).
fof(f83, plain, ((((e3 = unit) | ~ (e0 = op(e2, e3))) & (e3 = op(e2, e0))) | ~ sP76), inference(usedef, [], [e83])).
fof(e83, plain, (sP76 <=> (((e3 = unit) | ~ (e0 = op(e2, e3))) & (e3 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP76])])).
fof(f84, plain, ((((e3 = unit) | ~ (e1 = op(e2, e3))) & (e3 = op(e2, e1))) | ~ sP77), inference(usedef, [], [e84])).
fof(e84, plain, (sP77 <=> (((e3 = unit) | ~ (e1 = op(e2, e3))) & (e3 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP77])])).
fof(f85, plain, ((((e3 = unit) | ~ (e2 = op(e2, e3))) & (e3 = op(e2, e2))) | ~ sP78), inference(usedef, [], [e85])).
fof(e85, plain, (sP78 <=> (((e3 = unit) | ~ (e2 = op(e2, e3))) & (e3 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP78])])).
fof(f86, plain, ((((e3 = unit) | ~ (e3 = op(e2, e3))) & (e3 = op(e2, e3))) | ~ sP79), inference(usedef, [], [e86])).
fof(e86, plain, (sP79 <=> (((e3 = unit) | ~ (e3 = op(e2, e3))) & (e3 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP79])])).
fof(f87, plain, ((((e3 = unit) | ~ (e4 = op(e2, e3))) & (e3 = op(e2, e4))) | ~ sP80), inference(usedef, [], [e87])).
fof(e87, plain, (sP80 <=> (((e3 = unit) | ~ (e4 = op(e2, e3))) & (e3 = op(e2, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP80])])).
fof(f88, plain, ((((e3 = unit) | ~ (e0 = op(e3, e3))) & (e3 = op(e3, e0))) | ~ sP81), inference(usedef, [], [e88])).
fof(e88, plain, (sP81 <=> (((e3 = unit) | ~ (e0 = op(e3, e3))) & (e3 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP81])])).
fof(f89, plain, ((((e3 = unit) | ~ (e1 = op(e3, e3))) & (e3 = op(e3, e1))) | ~ sP82), inference(usedef, [], [e89])).
fof(e89, plain, (sP82 <=> (((e3 = unit) | ~ (e1 = op(e3, e3))) & (e3 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP82])])).
fof(f90, plain, ((((e3 = unit) | ~ (e2 = op(e3, e3))) & (e3 = op(e3, e2))) | ~ sP83), inference(usedef, [], [e90])).
fof(e90, plain, (sP83 <=> (((e3 = unit) | ~ (e2 = op(e3, e3))) & (e3 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP83])])).
fof(f91, plain, ((((e3 = unit) | ~ (e3 = op(e3, e3))) & (e3 = op(e3, e3))) | ~ sP84), inference(usedef, [], [e91])).
fof(e91, plain, (sP84 <=> (((e3 = unit) | ~ (e3 = op(e3, e3))) & (e3 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP84])])).
fof(f92, plain, ((((e3 = unit) | ~ (e4 = op(e3, e3))) & (e3 = op(e3, e4))) | ~ sP85), inference(usedef, [], [e92])).
fof(e92, plain, (sP85 <=> (((e3 = unit) | ~ (e4 = op(e3, e3))) & (e3 = op(e3, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP85])])).
fof(f93, plain, ((((e3 = unit) | ~ (e0 = op(e4, e3))) & (e3 = op(e4, e0))) | ~ sP86), inference(usedef, [], [e93])).
fof(e93, plain, (sP86 <=> (((e3 = unit) | ~ (e0 = op(e4, e3))) & (e3 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP86])])).
fof(f94, plain, ((((e3 = unit) | ~ (e1 = op(e4, e3))) & (e3 = op(e4, e1))) | ~ sP87), inference(usedef, [], [e94])).
fof(e94, plain, (sP87 <=> (((e3 = unit) | ~ (e1 = op(e4, e3))) & (e3 = op(e4, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP87])])).
fof(f95, plain, ((((e4 = unit) | ~ (e0 = op(e0, e4))) & (e4 = op(e0, e0))) | ~ sP88), inference(usedef, [], [e95])).
fof(e95, plain, (sP88 <=> (((e4 = unit) | ~ (e0 = op(e0, e4))) & (e4 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP88])])).
fof(f96, plain, ((((e4 = unit) | ~ (e1 = op(e0, e4))) & (e4 = op(e0, e1))) | ~ sP89), inference(usedef, [], [e96])).
fof(e96, plain, (sP89 <=> (((e4 = unit) | ~ (e1 = op(e0, e4))) & (e4 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP89])])).
fof(f97, plain, ((((e4 = unit) | ~ (e2 = op(e0, e4))) & (e4 = op(e0, e2))) | ~ sP90), inference(usedef, [], [e97])).
fof(e97, plain, (sP90 <=> (((e4 = unit) | ~ (e2 = op(e0, e4))) & (e4 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP90])])).
fof(f98, plain, ((((e4 = unit) | ~ (e3 = op(e0, e4))) & (e4 = op(e0, e3))) | ~ sP91), inference(usedef, [], [e98])).
fof(e98, plain, (sP91 <=> (((e4 = unit) | ~ (e3 = op(e0, e4))) & (e4 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP91])])).
fof(f99, plain, ((((e4 = unit) | ~ (e4 = op(e0, e4))) & (e4 = op(e0, e4))) | ~ sP92), inference(usedef, [], [e99])).
fof(e99, plain, (sP92 <=> (((e4 = unit) | ~ (e4 = op(e0, e4))) & (e4 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP92])])).
fof(f100, plain, ((((e4 = unit) | ~ (e0 = op(e1, e4))) & (e4 = op(e1, e0))) | ~ sP93), inference(usedef, [], [e100])).
fof(e100, plain, (sP93 <=> (((e4 = unit) | ~ (e0 = op(e1, e4))) & (e4 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP93])])).
fof(f101, plain, ((((e4 = unit) | ~ (e1 = op(e1, e4))) & (e4 = op(e1, e1))) | ~ sP94), inference(usedef, [], [e101])).
fof(e101, plain, (sP94 <=> (((e4 = unit) | ~ (e1 = op(e1, e4))) & (e4 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP94])])).
fof(f102, plain, ((((e4 = unit) | ~ (e2 = op(e1, e4))) & (e4 = op(e1, e2))) | ~ sP95), inference(usedef, [], [e102])).
fof(e102, plain, (sP95 <=> (((e4 = unit) | ~ (e2 = op(e1, e4))) & (e4 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP95])])).
fof(f103, plain, ((((e4 = unit) | ~ (e3 = op(e1, e4))) & (e4 = op(e1, e3))) | ~ sP96), inference(usedef, [], [e103])).
fof(e103, plain, (sP96 <=> (((e4 = unit) | ~ (e3 = op(e1, e4))) & (e4 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP96])])).
fof(f104, plain, ((((e4 = unit) | ~ (e4 = op(e1, e4))) & (e4 = op(e1, e4))) | ~ sP97), inference(usedef, [], [e104])).
fof(e104, plain, (sP97 <=> (((e4 = unit) | ~ (e4 = op(e1, e4))) & (e4 = op(e1, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP97])])).
fof(f105, plain, ((((e4 = unit) | ~ (e0 = op(e2, e4))) & (e4 = op(e2, e0))) | ~ sP98), inference(usedef, [], [e105])).
fof(e105, plain, (sP98 <=> (((e4 = unit) | ~ (e0 = op(e2, e4))) & (e4 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP98])])).
fof(f106, plain, ((((e4 = unit) | ~ (e1 = op(e2, e4))) & (e4 = op(e2, e1))) | ~ sP99), inference(usedef, [], [e106])).
fof(e106, plain, (sP99 <=> (((e4 = unit) | ~ (e1 = op(e2, e4))) & (e4 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP99])])).
fof(f107, plain, ((((e4 = unit) | ~ (e2 = op(e2, e4))) & (e4 = op(e2, e2))) | ~ sP100), inference(usedef, [], [e107])).
fof(e107, plain, (sP100 <=> (((e4 = unit) | ~ (e2 = op(e2, e4))) & (e4 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP100])])).
fof(f108, plain, ((((e4 = unit) | ~ (e3 = op(e2, e4))) & (e4 = op(e2, e3))) | ~ sP101), inference(usedef, [], [e108])).
fof(e108, plain, (sP101 <=> (((e4 = unit) | ~ (e3 = op(e2, e4))) & (e4 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP101])])).
fof(f109, plain, ((((e4 = unit) | ~ (e4 = op(e2, e4))) & (e4 = op(e2, e4))) | ~ sP102), inference(usedef, [], [e109])).
fof(e109, plain, (sP102 <=> (((e4 = unit) | ~ (e4 = op(e2, e4))) & (e4 = op(e2, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP102])])).
fof(f110, plain, ((((e4 = unit) | ~ (e0 = op(e3, e4))) & (e4 = op(e3, e0))) | ~ sP103), inference(usedef, [], [e110])).
fof(e110, plain, (sP103 <=> (((e4 = unit) | ~ (e0 = op(e3, e4))) & (e4 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP103])])).
fof(f111, plain, ((((e4 = unit) | ~ (e1 = op(e3, e4))) & (e4 = op(e3, e1))) | ~ sP104), inference(usedef, [], [e111])).
fof(e111, plain, (sP104 <=> (((e4 = unit) | ~ (e1 = op(e3, e4))) & (e4 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP104])])).
fof(f112, plain, ((((e4 = unit) | ~ (e2 = op(e3, e4))) & (e4 = op(e3, e2))) | ~ sP105), inference(usedef, [], [e112])).
fof(e112, plain, (sP105 <=> (((e4 = unit) | ~ (e2 = op(e3, e4))) & (e4 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP105])])).
fof(f113, plain, ((((e4 = unit) | ~ (e3 = op(e3, e4))) & (e4 = op(e3, e3))) | ~ sP106), inference(usedef, [], [e113])).
fof(e113, plain, (sP106 <=> (((e4 = unit) | ~ (e3 = op(e3, e4))) & (e4 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP106])])).
fof(f114, plain, ((((e4 = unit) | ~ (e4 = op(e3, e4))) & (e4 = op(e3, e4))) | ~ sP107), inference(usedef, [], [e114])).
fof(e114, plain, (sP107 <=> (((e4 = unit) | ~ (e4 = op(e3, e4))) & (e4 = op(e3, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP107])])).
fof(f115, plain, ((((e4 = unit) | ~ (e0 = op(e4, e4))) & (e4 = op(e4, e0))) | ~ sP108), inference(usedef, [], [e115])).
fof(e115, plain, (sP108 <=> (((e4 = unit) | ~ (e0 = op(e4, e4))) & (e4 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP108])])).
fof(f116, plain, ((((e4 = unit) | ~ (e1 = op(e4, e4))) & (e4 = op(e4, e1))) | ~ sP109), inference(usedef, [], [e116])).
fof(e116, plain, (sP109 <=> (((e4 = unit) | ~ (e1 = op(e4, e4))) & (e4 = op(e4, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP109])])).
fof(f117, plain, ((((e0 = unit) | ~ (e4 = op(e4, e0))) & (e0 = op(e4, e4))) | (((e0 = unit) | ~ (e3 = op(e4, e0))) & (e0 = op(e4, e3))) | (((e0 = unit) | ~ (e2 = op(e4, e0))) & (e0 = op(e4, e2))) | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ sP110), inference(usedef, [], [e117])).
fof(e117, plain, (sP110 <=> ((((e0 = unit) | ~ (e4 = op(e4, e0))) & (e0 = op(e4, e4))) | (((e0 = unit) | ~ (e3 = op(e4, e0))) & (e0 = op(e4, e3))) | (((e0 = unit) | ~ (e2 = op(e4, e0))) & (e0 = op(e4, e2))) | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0)), introduced(predicate_definition_introduction, [new_symbols(naming, [sP110])])).
fof(f118, plain, ((((e1 = unit) | ~ (e4 = op(e4, e1))) & (e1 = op(e4, e4))) | (((e1 = unit) | ~ (e3 = op(e4, e1))) & (e1 = op(e4, e3))) | (((e1 = unit) | ~ (e2 = op(e4, e1))) & (e1 = op(e4, e2))) | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | ~ sP111), inference(usedef, [], [e118])).
fof(e118, plain, (sP111 <=> ((((e1 = unit) | ~ (e4 = op(e4, e1))) & (e1 = op(e4, e4))) | (((e1 = unit) | ~ (e3 = op(e4, e1))) & (e1 = op(e4, e3))) | (((e1 = unit) | ~ (e2 = op(e4, e1))) & (e1 = op(e4, e2))) | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22)), introduced(predicate_definition_introduction, [new_symbols(naming, [sP111])])).
fof(f119, plain, ((((e2 = unit) | ~ (e4 = op(e4, e2))) & (e2 = op(e4, e4))) | (((e2 = unit) | ~ (e3 = op(e4, e2))) & (e2 = op(e4, e3))) | (((e2 = unit) | ~ (e2 = op(e4, e2))) & (e2 = op(e4, e2))) | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | ~ sP112), inference(usedef, [], [e119])).
fof(e119, plain, (sP112 <=> ((((e2 = unit) | ~ (e4 = op(e4, e2))) & (e2 = op(e4, e4))) | (((e2 = unit) | ~ (e3 = op(e4, e2))) & (e2 = op(e4, e3))) | (((e2 = unit) | ~ (e2 = op(e4, e2))) & (e2 = op(e4, e2))) | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44)), introduced(predicate_definition_introduction, [new_symbols(naming, [sP112])])).
fof(f120, plain, ((((e3 = unit) | ~ (e4 = op(e4, e3))) & (e3 = op(e4, e4))) | (((e3 = unit) | ~ (e3 = op(e4, e3))) & (e3 = op(e4, e3))) | (((e3 = unit) | ~ (e2 = op(e4, e3))) & (e3 = op(e4, e2))) | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | ~ sP113), inference(usedef, [], [e120])).
fof(e120, plain, (sP113 <=> ((((e3 = unit) | ~ (e4 = op(e4, e3))) & (e3 = op(e4, e4))) | (((e3 = unit) | ~ (e3 = op(e4, e3))) & (e3 = op(e4, e3))) | (((e3 = unit) | ~ (e2 = op(e4, e3))) & (e3 = op(e4, e2))) | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66)), introduced(predicate_definition_introduction, [new_symbols(naming, [sP113])])).
fof(f121, plain, ((((e4 = unit) | ~ (e4 = op(e4, e4))) & (e4 = op(e4, e4))) | (((e4 = unit) | ~ (e3 = op(e4, e4))) & (e4 = op(e4, e3))) | (((e4 = unit) | ~ (e2 = op(e4, e4))) & (e4 = op(e4, e2))) | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | ~ sP114), inference(usedef, [], [e121])).
fof(e121, plain, (sP114 <=> ((((e4 = unit) | ~ (e4 = op(e4, e4))) & (e4 = op(e4, e4))) | (((e4 = unit) | ~ (e3 = op(e4, e4))) & (e4 = op(e4, e3))) | (((e4 = unit) | ~ (e2 = op(e4, e4))) & (e4 = op(e4, e2))) | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88)), introduced(predicate_definition_introduction, [new_symbols(naming, [sP114])])).
fof(f122, plain, ((sP114 & sP113 & sP112 & sP111 & sP110) | ~ sP115), inference(usedef, [], [e122])).
fof(e122, plain, (sP115 <=> (sP114 & sP113 & sP112 & sP111 & sP110)), introduced(predicate_definition_introduction, [new_symbols(naming, [sP115])])).
fof(f6, plain, ((~ (e4 = op(e4, e4)) & ~ (e4 = op(e3, e4)) & ~ (e4 = op(e2, e4)) & ~ (e4 = op(e1, e4)) & ~ (e4 = op(e0, e4))) | (~ (e4 = op(e4, e4)) & ~ (e4 = op(e4, e3)) & ~ (e4 = op(e4, e2)) & ~ (e4 = op(e4, e1)) & ~ (e4 = op(e4, e0))) | (~ (e3 = op(e4, e4)) & ~ (e3 = op(e3, e4)) & ~ (e3 = op(e2, e4)) & ~ (e3 = op(e1, e4)) & ~ (e3 = op(e0, e4))) | (~ (e3 = op(e4, e4)) & ~ (e3 = op(e4, e3)) & ~ (e3 = op(e4, e2)) & ~ (e3 = op(e4, e1)) & ~ (e3 = op(e4, e0))) | (~ (e2 = op(e4, e4)) & ~ (e2 = op(e3, e4)) & ~ (e2 = op(e2, e4)) & ~ (e2 = op(e1, e4)) & ~ (e2 = op(e0, e4))) | (~ (e2 = op(e4, e4)) & ~ (e2 = op(e4, e3)) & ~ (e2 = op(e4, e2)) & ~ (e2 = op(e4, e1)) & ~ (e2 = op(e4, e0))) | (~ (e1 = op(e4, e4)) & ~ (e1 = op(e3, e4)) & ~ (e1 = op(e2, e4)) & ~ (e1 = op(e1, e4)) & ~ (e1 = op(e0, e4))) | (~ (e1 = op(e4, e4)) & ~ (e1 = op(e4, e3)) & ~ (e1 = op(e4, e2)) & ~ (e1 = op(e4, e1)) & ~ (e1 = op(e4, e0))) | (~ (e0 = op(e4, e4)) & ~ (e0 = op(e3, e4)) & ~ (e0 = op(e2, e4)) & ~ (e0 = op(e1, e4)) & ~ (e0 = op(e0, e4))) | (~ (e0 = op(e4, e4)) & ~ (e0 = op(e4, e3)) & ~ (e0 = op(e4, e2)) & ~ (e0 = op(e4, e1)) & ~ (e0 = op(e4, e0))) | (~ (e4 = op(e4, e3)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e0, e3))) | (~ (e4 = op(e3, e4)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e3, e0))) | (~ (e3 = op(e4, e3)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e0, e3))) | (~ (e3 = op(e3, e4)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e3, e0))) | (~ (e2 = op(e4, e3)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e0, e3))) | (~ (e2 = op(e3, e4)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e3, e0))) | (~ (e1 = op(e4, e3)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e0, e3))) | (~ (e1 = op(e3, e4)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e3, e0))) | (~ (e0 = op(e4, e3)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e0, e3))) | (~ (e0 = op(e3, e4)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e3, e0))) | (~ (e4 = op(e4, e2)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e0, e2))) | (~ (e4 = op(e2, e4)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e2, e0))) | (~ (e3 = op(e4, e2)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e0, e2))) | (~ (e3 = op(e2, e4)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e2, e0))) | (~ (e2 = op(e4, e2)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e0, e2))) | (~ (e2 = op(e2, e4)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e2, e0))) | (~ (e1 = op(e4, e2)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e0, e2))) | (~ (e1 = op(e2, e4)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e2, e0))) | (~ (e0 = op(e4, e2)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e0, e2))) | (~ (e0 = op(e2, e4)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e2, e0))) | (~ (e4 = op(e4, e1)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e0, e1))) | (~ (e4 = op(e1, e4)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e1, e0))) | (~ (e3 = op(e4, e1)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e0, e1))) | (~ (e3 = op(e1, e4)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e1, e0))) | (~ (e2 = op(e4, e1)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e0, e1))) | (~ (e2 = op(e1, e4)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e1, e0))) | (~ (e1 = op(e4, e1)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e0, e1))) | (~ (e1 = op(e1, e4)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e1, e0))) | (~ (e0 = op(e4, e1)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e0, e1))) | (~ (e0 = op(e1, e4)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e1, e0))) | (~ (e4 = op(e4, e0)) & ~ (e4 = op(e3, e0)) & ~ (e4 = op(e2, e0)) & ~ (e4 = op(e1, e0)) & ~ (e4 = op(e0, e0))) | (~ (e4 = op(e0, e4)) & ~ (e4 = op(e0, e3)) & ~ (e4 = op(e0, e2)) & ~ (e4 = op(e0, e1)) & ~ (e4 = op(e0, e0))) | (~ (e3 = op(e4, e0)) & ~ (e3 = op(e3, e0)) & ~ (e3 = op(e2, e0)) & ~ (e3 = op(e1, e0)) & ~ (e3 = op(e0, e0))) | (~ (e3 = op(e0, e4)) & ~ (e3 = op(e0, e3)) & ~ (e3 = op(e0, e2)) & ~ (e3 = op(e0, e1)) & ~ (e3 = op(e0, e0))) | (~ (e2 = op(e4, e0)) & ~ (e2 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e2 = op(e1, e0)) & ~ (e2 = op(e0, e0))) | (~ (e2 = op(e0, e4)) & ~ (e2 = op(e0, e3)) & ~ (e2 = op(e0, e2)) & ~ (e2 = op(e0, e1)) & ~ (e2 = op(e0, e0))) | (~ (e1 = op(e4, e0)) & ~ (e1 = op(e3, e0)) & ~ (e1 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e1 = op(e0, e0))) | (~ (e1 = op(e0, e4)) & ~ (e1 = op(e0, e3)) & ~ (e1 = op(e0, e2)) & ~ (e1 = op(e0, e1)) & ~ (e1 = op(e0, e0))) | (~ (e0 = op(e4, e0)) & ~ (e0 = op(e3, e0)) & ~ (e0 = op(e2, e0)) & ~ (e0 = op(e1, e0)) & ~ (e0 = op(e0, e0))) | (~ (e0 = op(e0, e4)) & ~ (e0 = op(e0, e3)) & ~ (e0 = op(e0, e2)) & ~ (e0 = op(e0, e1)) & ~ (e0 = op(e0, e0))) | (~ (e4 = unit) & ~ (e3 = unit) & ~ (e2 = unit) & ~ (e1 = unit) & ~ (e0 = unit)) | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (e0 = op(e0, unit)) | ~ (e0 = op(unit, e0)) | (~ (e4 = op(e4, e4)) & ~ (e3 = op(e4, e4)) & ~ (e2 = op(e4, e4)) & ~ (e1 = op(e4, e4)) & ~ (e0 = op(e4, e4))) | (~ (e4 = op(e4, e3)) & ~ (e3 = op(e4, e3)) & ~ (e2 = op(e4, e3)) & ~ (e1 = op(e4, e3)) & ~ (e0 = op(e4, e3))) | (~ (e4 = op(e4, e2)) & ~ (e3 = op(e4, e2)) & ~ (e2 = op(e4, e2)) & ~ (e1 = op(e4, e2)) & ~ (e0 = op(e4, e2))) | (~ (e4 = op(e4, e1)) & ~ (e3 = op(e4, e1)) & ~ (e2 = op(e4, e1)) & ~ (e1 = op(e4, e1)) & ~ (e0 = op(e4, e1))) | (~ (e4 = op(e4, e0)) & ~ (e3 = op(e4, e0)) & ~ (e2 = op(e4, e0)) & ~ (e1 = op(e4, e0)) & ~ (e0 = op(e4, e0))) | (~ (e4 = op(e3, e4)) & ~ (e3 = op(e3, e4)) & ~ (e2 = op(e3, e4)) & ~ (e1 = op(e3, e4)) & ~ (e0 = op(e3, e4))) | (~ (e4 = op(e3, e3)) & ~ (e3 = op(e3, e3)) & ~ (e2 = op(e3, e3)) & ~ (e1 = op(e3, e3)) & ~ (e0 = op(e3, e3))) | (~ (e4 = op(e3, e2)) & ~ (e3 = op(e3, e2)) & ~ (e2 = op(e3, e2)) & ~ (e1 = op(e3, e2)) & ~ (e0 = op(e3, e2))) | (~ (e4 = op(e3, e1)) & ~ (e3 = op(e3, e1)) & ~ (e2 = op(e3, e1)) & ~ (e1 = op(e3, e1)) & ~ (e0 = op(e3, e1))) | (~ (e4 = op(e3, e0)) & ~ (e3 = op(e3, e0)) & ~ (e2 = op(e3, e0)) & ~ (e1 = op(e3, e0)) & ~ (e0 = op(e3, e0))) | (~ (e4 = op(e2, e4)) & ~ (e3 = op(e2, e4)) & ~ (e2 = op(e2, e4)) & ~ (e1 = op(e2, e4)) & ~ (e0 = op(e2, e4))) | (~ (e4 = op(e2, e3)) & ~ (e3 = op(e2, e3)) & ~ (e2 = op(e2, e3)) & ~ (e1 = op(e2, e3)) & ~ (e0 = op(e2, e3))) | (~ (e4 = op(e2, e2)) & ~ (e3 = op(e2, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e2, e2)) & ~ (e0 = op(e2, e2))) | (~ (e4 = op(e2, e1)) & ~ (e3 = op(e2, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e2, e1)) & ~ (e0 = op(e2, e1))) | (~ (e4 = op(e2, e0)) & ~ (e3 = op(e2, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e2, e0)) & ~ (e0 = op(e2, e0))) | (~ (e4 = op(e1, e4)) & ~ (e3 = op(e1, e4)) & ~ (e2 = op(e1, e4)) & ~ (e1 = op(e1, e4)) & ~ (e0 = op(e1, e4))) | (~ (e4 = op(e1, e3)) & ~ (e3 = op(e1, e3)) & ~ (e2 = op(e1, e3)) & ~ (e1 = op(e1, e3)) & ~ (e0 = op(e1, e3))) | (~ (e4 = op(e1, e2)) & ~ (e3 = op(e1, e2)) & ~ (e2 = op(e1, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e1, e2))) | (~ (e4 = op(e1, e1)) & ~ (e3 = op(e1, e1)) & ~ (e2 = op(e1, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e1, e1))) | (~ (e4 = op(e1, e0)) & ~ (e3 = op(e1, e0)) & ~ (e2 = op(e1, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e1, e0))) | (~ (e4 = op(e0, e4)) & ~ (e3 = op(e0, e4)) & ~ (e2 = op(e0, e4)) & ~ (e1 = op(e0, e4)) & ~ (e0 = op(e0, e4))) | (~ (e4 = op(e0, e3)) & ~ (e3 = op(e0, e3)) & ~ (e2 = op(e0, e3)) & ~ (e1 = op(e0, e3)) & ~ (e0 = op(e0, e3))) | (~ (e4 = op(e0, e2)) & ~ (e3 = op(e0, e2)) & ~ (e2 = op(e0, e2)) & ~ (e1 = op(e0, e2)) & ~ (e0 = op(e0, e2))) | (~ (e4 = op(e0, e1)) & ~ (e3 = op(e0, e1)) & ~ (e2 = op(e0, e1)) & ~ (e1 = op(e0, e1)) & ~ (e0 = op(e0, e1))) | (~ (e4 = op(e0, e0)) & ~ (e3 = op(e0, e0)) & ~ (e2 = op(e0, e0)) & ~ (e1 = op(e0, e0)) & ~ (e0 = op(e0, e0))) | (((((e4 = unit) | ~ (e4 = op(e4, e4))) & (e4 = op(e4, e4))) | (((e4 = unit) | ~ (e3 = op(e4, e4))) & (e4 = op(e4, e3))) | (((e4 = unit) | ~ (e2 = op(e4, e4))) & (e4 = op(e4, e2))) | (((e4 = unit) | ~ (e1 = op(e4, e4))) & (e4 = op(e4, e1))) | (((e4 = unit) | ~ (e0 = op(e4, e4))) & (e4 = op(e4, e0))) | (((e4 = unit) | ~ (e4 = op(e3, e4))) & (e4 = op(e3, e4))) | (((e4 = unit) | ~ (e3 = op(e3, e4))) & (e4 = op(e3, e3))) | (((e4 = unit) | ~ (e2 = op(e3, e4))) & (e4 = op(e3, e2))) | (((e4 = unit) | ~ (e1 = op(e3, e4))) & (e4 = op(e3, e1))) | (((e4 = unit) | ~ (e0 = op(e3, e4))) & (e4 = op(e3, e0))) | (((e4 = unit) | ~ (e4 = op(e2, e4))) & (e4 = op(e2, e4))) | (((e4 = unit) | ~ (e3 = op(e2, e4))) & (e4 = op(e2, e3))) | (((e4 = unit) | ~ (e2 = op(e2, e4))) & (e4 = op(e2, e2))) | (((e4 = unit) | ~ (e1 = op(e2, e4))) & (e4 = op(e2, e1))) | (((e4 = unit) | ~ (e0 = op(e2, e4))) & (e4 = op(e2, e0))) | (((e4 = unit) | ~ (e4 = op(e1, e4))) & (e4 = op(e1, e4))) | (((e4 = unit) | ~ (e3 = op(e1, e4))) & (e4 = op(e1, e3))) | (((e4 = unit) | ~ (e2 = op(e1, e4))) & (e4 = op(e1, e2))) | (((e4 = unit) | ~ (e1 = op(e1, e4))) & (e4 = op(e1, e1))) | (((e4 = unit) | ~ (e0 = op(e1, e4))) & (e4 = op(e1, e0))) | (((e4 = unit) | ~ (e4 = op(e0, e4))) & (e4 = op(e0, e4))) | (((e4 = unit) | ~ (e3 = op(e0, e4))) & (e4 = op(e0, e3))) | (((e4 = unit) | ~ (e2 = op(e0, e4))) & (e4 = op(e0, e2))) | (((e4 = unit) | ~ (e1 = op(e0, e4))) & (e4 = op(e0, e1))) | (((e4 = unit) | ~ (e0 = op(e0, e4))) & (e4 = op(e0, e0)))) & ((((e3 = unit) | ~ (e4 = op(e4, e3))) & (e3 = op(e4, e4))) | (((e3 = unit) | ~ (e3 = op(e4, e3))) & (e3 = op(e4, e3))) | (((e3 = unit) | ~ (e2 = op(e4, e3))) & (e3 = op(e4, e2))) | (((e3 = unit) | ~ (e1 = op(e4, e3))) & (e3 = op(e4, e1))) | (((e3 = unit) | ~ (e0 = op(e4, e3))) & (e3 = op(e4, e0))) | (((e3 = unit) | ~ (e4 = op(e3, e3))) & (e3 = op(e3, e4))) | (((e3 = unit) | ~ (e3 = op(e3, e3))) & (e3 = op(e3, e3))) | (((e3 = unit) | ~ (e2 = op(e3, e3))) & (e3 = op(e3, e2))) | (((e3 = unit) | ~ (e1 = op(e3, e3))) & (e3 = op(e3, e1))) | (((e3 = unit) | ~ (e0 = op(e3, e3))) & (e3 = op(e3, e0))) | (((e3 = unit) | ~ (e4 = op(e2, e3))) & (e3 = op(e2, e4))) | (((e3 = unit) | ~ (e3 = op(e2, e3))) & (e3 = op(e2, e3))) | (((e3 = unit) | ~ (e2 = op(e2, e3))) & (e3 = op(e2, e2))) | (((e3 = unit) | ~ (e1 = op(e2, e3))) & (e3 = op(e2, e1))) | (((e3 = unit) | ~ (e0 = op(e2, e3))) & (e3 = op(e2, e0))) | (((e3 = unit) | ~ (e4 = op(e1, e3))) & (e3 = op(e1, e4))) | (((e3 = unit) | ~ (e3 = op(e1, e3))) & (e3 = op(e1, e3))) | (((e3 = unit) | ~ (e2 = op(e1, e3))) & (e3 = op(e1, e2))) | (((e3 = unit) | ~ (e1 = op(e1, e3))) & (e3 = op(e1, e1))) | (((e3 = unit) | ~ (e0 = op(e1, e3))) & (e3 = op(e1, e0))) | (((e3 = unit) | ~ (e4 = op(e0, e3))) & (e3 = op(e0, e4))) | (((e3 = unit) | ~ (e3 = op(e0, e3))) & (e3 = op(e0, e3))) | (((e3 = unit) | ~ (e2 = op(e0, e3))) & (e3 = op(e0, e2))) | (((e3 = unit) | ~ (e1 = op(e0, e3))) & (e3 = op(e0, e1))) | (((e3 = unit) | ~ (e0 = op(e0, e3))) & (e3 = op(e0, e0)))) & ((((e2 = unit) | ~ (e4 = op(e4, e2))) & (e2 = op(e4, e4))) | (((e2 = unit) | ~ (e3 = op(e4, e2))) & (e2 = op(e4, e3))) | (((e2 = unit) | ~ (e2 = op(e4, e2))) & (e2 = op(e4, e2))) | (((e2 = unit) | ~ (e1 = op(e4, e2))) & (e2 = op(e4, e1))) | (((e2 = unit) | ~ (e0 = op(e4, e2))) & (e2 = op(e4, e0))) | (((e2 = unit) | ~ (e4 = op(e3, e2))) & (e2 = op(e3, e4))) | (((e2 = unit) | ~ (e3 = op(e3, e2))) & (e2 = op(e3, e3))) | (((e2 = unit) | ~ (e2 = op(e3, e2))) & (e2 = op(e3, e2))) | (((e2 = unit) | ~ (e1 = op(e3, e2))) & (e2 = op(e3, e1))) | (((e2 = unit) | ~ (e0 = op(e3, e2))) & (e2 = op(e3, e0))) | (((e2 = unit) | ~ (e4 = op(e2, e2))) & (e2 = op(e2, e4))) | (((e2 = unit) | ~ (e3 = op(e2, e2))) & (e2 = op(e2, e3))) | (((e2 = unit) | ~ (e2 = op(e2, e2))) & (e2 = op(e2, e2))) | (((e2 = unit) | ~ (e1 = op(e2, e2))) & (e2 = op(e2, e1))) | (((e2 = unit) | ~ (e0 = op(e2, e2))) & (e2 = op(e2, e0))) | (((e2 = unit) | ~ (e4 = op(e1, e2))) & (e2 = op(e1, e4))) | (((e2 = unit) | ~ (e3 = op(e1, e2))) & (e2 = op(e1, e3))) | (((e2 = unit) | ~ (e2 = op(e1, e2))) & (e2 = op(e1, e2))) | (((e2 = unit) | ~ (e1 = op(e1, e2))) & (e2 = op(e1, e1))) | (((e2 = unit) | ~ (e0 = op(e1, e2))) & (e2 = op(e1, e0))) | (((e2 = unit) | ~ (e4 = op(e0, e2))) & (e2 = op(e0, e4))) | (((e2 = unit) | ~ (e3 = op(e0, e2))) & (e2 = op(e0, e3))) | (((e2 = unit) | ~ (e2 = op(e0, e2))) & (e2 = op(e0, e2))) | (((e2 = unit) | ~ (e1 = op(e0, e2))) & (e2 = op(e0, e1))) | (((e2 = unit) | ~ (e0 = op(e0, e2))) & (e2 = op(e0, e0)))) & ((((e1 = unit) | ~ (e4 = op(e4, e1))) & (e1 = op(e4, e4))) | (((e1 = unit) | ~ (e3 = op(e4, e1))) & (e1 = op(e4, e3))) | (((e1 = unit) | ~ (e2 = op(e4, e1))) & (e1 = op(e4, e2))) | (((e1 = unit) | ~ (e1 = op(e4, e1))) & (e1 = op(e4, e1))) | (((e1 = unit) | ~ (e0 = op(e4, e1))) & (e1 = op(e4, e0))) | (((e1 = unit) | ~ (e4 = op(e3, e1))) & (e1 = op(e3, e4))) | (((e1 = unit) | ~ (e3 = op(e3, e1))) & (e1 = op(e3, e3))) | (((e1 = unit) | ~ (e2 = op(e3, e1))) & (e1 = op(e3, e2))) | (((e1 = unit) | ~ (e1 = op(e3, e1))) & (e1 = op(e3, e1))) | (((e1 = unit) | ~ (e0 = op(e3, e1))) & (e1 = op(e3, e0))) | (((e1 = unit) | ~ (e4 = op(e2, e1))) & (e1 = op(e2, e4))) | (((e1 = unit) | ~ (e3 = op(e2, e1))) & (e1 = op(e2, e3))) | (((e1 = unit) | ~ (e2 = op(e2, e1))) & (e1 = op(e2, e2))) | (((e1 = unit) | ~ (e1 = op(e2, e1))) & (e1 = op(e2, e1))) | (((e1 = unit) | ~ (e0 = op(e2, e1))) & (e1 = op(e2, e0))) | (((e1 = unit) | ~ (e4 = op(e1, e1))) & (e1 = op(e1, e4))) | (((e1 = unit) | ~ (e3 = op(e1, e1))) & (e1 = op(e1, e3))) | (((e1 = unit) | ~ (e2 = op(e1, e1))) & (e1 = op(e1, e2))) | (((e1 = unit) | ~ (e1 = op(e1, e1))) & (e1 = op(e1, e1))) | (((e1 = unit) | ~ (e0 = op(e1, e1))) & (e1 = op(e1, e0))) | (((e1 = unit) | ~ (e4 = op(e0, e1))) & (e1 = op(e0, e4))) | (((e1 = unit) | ~ (e3 = op(e0, e1))) & (e1 = op(e0, e3))) | (((e1 = unit) | ~ (e2 = op(e0, e1))) & (e1 = op(e0, e2))) | (((e1 = unit) | ~ (e1 = op(e0, e1))) & (e1 = op(e0, e1))) | (((e1 = unit) | ~ (e0 = op(e0, e1))) & (e1 = op(e0, e0)))) & ((((e0 = unit) | ~ (e4 = op(e4, e0))) & (e0 = op(e4, e4))) | (((e0 = unit) | ~ (e3 = op(e4, e0))) & (e0 = op(e4, e3))) | (((e0 = unit) | ~ (e2 = op(e4, e0))) & (e0 = op(e4, e2))) | (((e0 = unit) | ~ (e1 = op(e4, e0))) & (e0 = op(e4, e1))) | (((e0 = unit) | ~ (e0 = op(e4, e0))) & (e0 = op(e4, e0))) | (((e0 = unit) | ~ (e4 = op(e3, e0))) & (e0 = op(e3, e4))) | (((e0 = unit) | ~ (e3 = op(e3, e0))) & (e0 = op(e3, e3))) | (((e0 = unit) | ~ (e2 = op(e3, e0))) & (e0 = op(e3, e2))) | (((e0 = unit) | ~ (e1 = op(e3, e0))) & (e0 = op(e3, e1))) | (((e0 = unit) | ~ (e0 = op(e3, e0))) & (e0 = op(e3, e0))) | (((e0 = unit) | ~ (e4 = op(e2, e0))) & (e0 = op(e2, e4))) | (((e0 = unit) | ~ (e3 = op(e2, e0))) & (e0 = op(e2, e3))) | (((e0 = unit) | ~ (e2 = op(e2, e0))) & (e0 = op(e2, e2))) | (((e0 = unit) | ~ (e1 = op(e2, e0))) & (e0 = op(e2, e1))) | (((e0 = unit) | ~ (e0 = op(e2, e0))) & (e0 = op(e2, e0))) | (((e0 = unit) | ~ (e4 = op(e1, e0))) & (e0 = op(e1, e4))) | (((e0 = unit) | ~ (e3 = op(e1, e0))) & (e0 = op(e1, e3))) | (((e0 = unit) | ~ (e2 = op(e1, e0))) & (e0 = op(e1, e2))) | (((e0 = unit) | ~ (e1 = op(e1, e0))) & (e0 = op(e1, e1))) | (((e0 = unit) | ~ (e0 = op(e1, e0))) & (e0 = op(e1, e0))) | (((e0 = unit) | ~ (e4 = op(e0, e0))) & (e0 = op(e0, e4))) | (((e0 = unit) | ~ (e3 = op(e0, e0))) & (e0 = op(e0, e3))) | (((e0 = unit) | ~ (e2 = op(e0, e0))) & (e0 = op(e0, e2))) | (((e0 = unit) | ~ (e1 = op(e0, e0))) & (e0 = op(e0, e1))) | (((e0 = unit) | ~ (e0 = op(e0, e0))) & (e0 = op(e0, e0)))))), inference(ennf_transformation, [], [f5])).
fof(f5, plain, ~ (((e4 = op(e4, e4)) | (e4 = op(e3, e4)) | (e4 = op(e2, e4)) | (e4 = op(e1, e4)) | (e4 = op(e0, e4))) & ((e4 = op(e4, e4)) | (e4 = op(e4, e3)) | (e4 = op(e4, e2)) | (e4 = op(e4, e1)) | (e4 = op(e4, e0))) & ((e3 = op(e4, e4)) | (e3 = op(e3, e4)) | (e3 = op(e2, e4)) | (e3 = op(e1, e4)) | (e3 = op(e0, e4))) & ((e3 = op(e4, e4)) | (e3 = op(e4, e3)) | (e3 = op(e4, e2)) | (e3 = op(e4, e1)) | (e3 = op(e4, e0))) & ((e2 = op(e4, e4)) | (e2 = op(e3, e4)) | (e2 = op(e2, e4)) | (e2 = op(e1, e4)) | (e2 = op(e0, e4))) & ((e2 = op(e4, e4)) | (e2 = op(e4, e3)) | (e2 = op(e4, e2)) | (e2 = op(e4, e1)) | (e2 = op(e4, e0))) & ((e1 = op(e4, e4)) | (e1 = op(e3, e4)) | (e1 = op(e2, e4)) | (e1 = op(e1, e4)) | (e1 = op(e0, e4))) & ((e1 = op(e4, e4)) | (e1 = op(e4, e3)) | (e1 = op(e4, e2)) | (e1 = op(e4, e1)) | (e1 = op(e4, e0))) & ((e0 = op(e4, e4)) | (e0 = op(e3, e4)) | (e0 = op(e2, e4)) | (e0 = op(e1, e4)) | (e0 = op(e0, e4))) & ((e0 = op(e4, e4)) | (e0 = op(e4, e3)) | (e0 = op(e4, e2)) | (e0 = op(e4, e1)) | (e0 = op(e4, e0))) & ((e4 = op(e4, e3)) | (e4 = op(e3, e3)) | (e4 = op(e2, e3)) | (e4 = op(e1, e3)) | (e4 = op(e0, e3))) & ((e4 = op(e3, e4)) | (e4 = op(e3, e3)) | (e4 = op(e3, e2)) | (e4 = op(e3, e1)) | (e4 = op(e3, e0))) & ((e3 = op(e4, e3)) | (e3 = op(e3, e3)) | (e3 = op(e2, e3)) | (e3 = op(e1, e3)) | (e3 = op(e0, e3))) & ((e3 = op(e3, e4)) | (e3 = op(e3, e3)) | (e3 = op(e3, e2)) | (e3 = op(e3, e1)) | (e3 = op(e3, e0))) & ((e2 = op(e4, e3)) | (e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(e0, e3))) & ((e2 = op(e3, e4)) | (e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, e0))) & ((e1 = op(e4, e3)) | (e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(e0, e3))) & ((e1 = op(e3, e4)) | (e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, e0))) & ((e0 = op(e4, e3)) | (e0 = op(e3, e3)) | (e0 = op(e2, e3)) | (e0 = op(e1, e3)) | (e0 = op(e0, e3))) & ((e0 = op(e3, e4)) | (e0 = op(e3, e3)) | (e0 = op(e3, e2)) | (e0 = op(e3, e1)) | (e0 = op(e3, e0))) & ((e4 = op(e4, e2)) | (e4 = op(e3, e2)) | (e4 = op(e2, e2)) | (e4 = op(e1, e2)) | (e4 = op(e0, e2))) & ((e4 = op(e2, e4)) | (e4 = op(e2, e3)) | (e4 = op(e2, e2)) | (e4 = op(e2, e1)) | (e4 = op(e2, e0))) & ((e3 = op(e4, e2)) | (e3 = op(e3, e2)) | (e3 = op(e2, e2)) | (e3 = op(e1, e2)) | (e3 = op(e0, e2))) & ((e3 = op(e2, e4)) | (e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, e0))) & ((e2 = op(e4, e2)) | (e2 = op(e3, e2)) | (e2 = op(e2, e2)) | (e2 = op(e1, e2)) | (e2 = op(e0, e2))) & ((e2 = op(e2, e4)) | (e2 = op(e2, e3)) | (e2 = op(e2, e2)) | (e2 = op(e2, e1)) | (e2 = op(e2, e0))) & ((e1 = op(e4, e2)) | (e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(e0, e2))) & ((e1 = op(e2, e4)) | (e1 = op(e2, e3)) | (e1 = op(e2, e2)) | (e1 = op(e2, e1)) | (e1 = op(e2, e0))) & ((e0 = op(e4, e2)) | (e0 = op(e3, e2)) | (e0 = op(e2, e2)) | (e0 = op(e1, e2)) | (e0 = op(e0, e2))) & ((e0 = op(e2, e4)) | (e0 = op(e2, e3)) | (e0 = op(e2, e2)) | (e0 = op(e2, e1)) | (e0 = op(e2, e0))) & ((e4 = op(e4, e1)) | (e4 = op(e3, e1)) | (e4 = op(e2, e1)) | (e4 = op(e1, e1)) | (e4 = op(e0, e1))) & ((e4 = op(e1, e4)) | (e4 = op(e1, e3)) | (e4 = op(e1, e2)) | (e4 = op(e1, e1)) | (e4 = op(e1, e0))) & ((e3 = op(e4, e1)) | (e3 = op(e3, e1)) | (e3 = op(e2, e1)) | (e3 = op(e1, e1)) | (e3 = op(e0, e1))) & ((e3 = op(e1, e4)) | (e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, e0))) & ((e2 = op(e4, e1)) | (e2 = op(e3, e1)) | (e2 = op(e2, e1)) | (e2 = op(e1, e1)) | (e2 = op(e0, e1))) & ((e2 = op(e1, e4)) | (e2 = op(e1, e3)) | (e2 = op(e1, e2)) | (e2 = op(e1, e1)) | (e2 = op(e1, e0))) & ((e1 = op(e4, e1)) | (e1 = op(e3, e1)) | (e1 = op(e2, e1)) | (e1 = op(e1, e1)) | (e1 = op(e0, e1))) & ((e1 = op(e1, e4)) | (e1 = op(e1, e3)) | (e1 = op(e1, e2)) | (e1 = op(e1, e1)) | (e1 = op(e1, e0))) & ((e0 = op(e4, e1)) | (e0 = op(e3, e1)) | (e0 = op(e2, e1)) | (e0 = op(e1, e1)) | (e0 = op(e0, e1))) & ((e0 = op(e1, e4)) | (e0 = op(e1, e3)) | (e0 = op(e1, e2)) | (e0 = op(e1, e1)) | (e0 = op(e1, e0))) & ((e4 = op(e4, e0)) | (e4 = op(e3, e0)) | (e4 = op(e2, e0)) | (e4 = op(e1, e0)) | (e4 = op(e0, e0))) & ((e4 = op(e0, e4)) | (e4 = op(e0, e3)) | (e4 = op(e0, e2)) | (e4 = op(e0, e1)) | (e4 = op(e0, e0))) & ((e3 = op(e4, e0)) | (e3 = op(e3, e0)) | (e3 = op(e2, e0)) | (e3 = op(e1, e0)) | (e3 = op(e0, e0))) & ((e3 = op(e0, e4)) | (e3 = op(e0, e3)) | (e3 = op(e0, e2)) | (e3 = op(e0, e1)) | (e3 = op(e0, e0))) & ((e2 = op(e4, e0)) | (e2 = op(e3, e0)) | (e2 = op(e2, e0)) | (e2 = op(e1, e0)) | (e2 = op(e0, e0))) & ((e2 = op(e0, e4)) | (e2 = op(e0, e3)) | (e2 = op(e0, e2)) | (e2 = op(e0, e1)) | (e2 = op(e0, e0))) & ((e1 = op(e4, e0)) | (e1 = op(e3, e0)) | (e1 = op(e2, e0)) | (e1 = op(e1, e0)) | (e1 = op(e0, e0))) & ((e1 = op(e0, e4)) | (e1 = op(e0, e3)) | (e1 = op(e0, e2)) | (e1 = op(e0, e1)) | (e1 = op(e0, e0))) & ((e0 = op(e4, e0)) | (e0 = op(e3, e0)) | (e0 = op(e2, e0)) | (e0 = op(e1, e0)) | (e0 = op(e0, e0))) & ((e0 = op(e0, e4)) | (e0 = op(e0, e3)) | (e0 = op(e0, e2)) | (e0 = op(e0, e1)) | (e0 = op(e0, e0))) & ((e4 = unit) | (e3 = unit) | (e2 = unit) | (e1 = unit) | (e0 = unit)) & (e4 = op(e4, unit)) & (e4 = op(unit, e4)) & (e3 = op(e3, unit)) & (e3 = op(unit, e3)) & (e2 = op(e2, unit)) & (e2 = op(unit, e2)) & (e1 = op(e1, unit)) & (e1 = op(unit, e1)) & (e0 = op(e0, unit)) & (e0 = op(unit, e0)) & ((e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | (e1 = op(e4, e4)) | (e0 = op(e4, e4))) & ((e4 = op(e4, e3)) | (e3 = op(e4, e3)) | (e2 = op(e4, e3)) | (e1 = op(e4, e3)) | (e0 = op(e4, e3))) & ((e4 = op(e4, e2)) | (e3 = op(e4, e2)) | (e2 = op(e4, e2)) | (e1 = op(e4, e2)) | (e0 = op(e4, e2))) & ((e4 = op(e4, e1)) | (e3 = op(e4, e1)) | (e2 = op(e4, e1)) | (e1 = op(e4, e1)) | (e0 = op(e4, e1))) & ((e4 = op(e4, e0)) | (e3 = op(e4, e0)) | (e2 = op(e4, e0)) | (e1 = op(e4, e0)) | (e0 = op(e4, e0))) & ((e4 = op(e3, e4)) | (e3 = op(e3, e4)) | (e2 = op(e3, e4)) | (e1 = op(e3, e4)) | (e0 = op(e3, e4))) & ((e4 = op(e3, e3)) | (e3 = op(e3, e3)) | (e2 = op(e3, e3)) | (e1 = op(e3, e3)) | (e0 = op(e3, e3))) & ((e4 = op(e3, e2)) | (e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (e0 = op(e3, e2))) & ((e4 = op(e3, e1)) | (e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (e0 = op(e3, e1))) & ((e4 = op(e3, e0)) | (e3 = op(e3, e0)) | (e2 = op(e3, e0)) | (e1 = op(e3, e0)) | (e0 = op(e3, e0))) & ((e4 = op(e2, e4)) | (e3 = op(e2, e4)) | (e2 = op(e2, e4)) | (e1 = op(e2, e4)) | (e0 = op(e2, e4))) & ((e4 = op(e2, e3)) | (e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))) & ((e4 = op(e2, e2)) | (e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (e0 = op(e2, e2))) & ((e4 = op(e2, e1)) | (e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))) & ((e4 = op(e2, e0)) | (e3 = op(e2, e0)) | (e2 = op(e2, e0)) | (e1 = op(e2, e0)) | (e0 = op(e2, e0))) & ((e4 = op(e1, e4)) | (e3 = op(e1, e4)) | (e2 = op(e1, e4)) | (e1 = op(e1, e4)) | (e0 = op(e1, e4))) & ((e4 = op(e1, e3)) | (e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (e0 = op(e1, e3))) & ((e4 = op(e1, e2)) | (e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (e0 = op(e1, e2))) & ((e4 = op(e1, e1)) | (e3 = op(e1, e1)) | (e2 = op(e1, e1)) | (e1 = op(e1, e1)) | (e0 = op(e1, e1))) & ((e4 = op(e1, e0)) | (e3 = op(e1, e0)) | (e2 = op(e1, e0)) | (e1 = op(e1, e0)) | (e0 = op(e1, e0))) & ((e4 = op(e0, e4)) | (e3 = op(e0, e4)) | (e2 = op(e0, e4)) | (e1 = op(e0, e4)) | (e0 = op(e0, e4))) & ((e4 = op(e0, e3)) | (e3 = op(e0, e3)) | (e2 = op(e0, e3)) | (e1 = op(e0, e3)) | (e0 = op(e0, e3))) & ((e4 = op(e0, e2)) | (e3 = op(e0, e2)) | (e2 = op(e0, e2)) | (e1 = op(e0, e2)) | (e0 = op(e0, e2))) & ((e4 = op(e0, e1)) | (e3 = op(e0, e1)) | (e2 = op(e0, e1)) | (e1 = op(e0, e1)) | (e0 = op(e0, e1))) & ((e4 = op(e0, e0)) | (e3 = op(e0, e0)) | (e2 = op(e0, e0)) | (e1 = op(e0, e0)) | (e0 = op(e0, e0))) & ((((~ (e4 = unit) & (e4 = op(e4, e4))) | ~ (e4 = op(e4, e4))) & ((~ (e4 = unit) & (e3 = op(e4, e4))) | ~ (e4 = op(e4, e3))) & ((~ (e4 = unit) & (e2 = op(e4, e4))) | ~ (e4 = op(e4, e2))) & ((~ (e4 = unit) & (e1 = op(e4, e4))) | ~ (e4 = op(e4, e1))) & ((~ (e4 = unit) & (e0 = op(e4, e4))) | ~ (e4 = op(e4, e0))) & ((~ (e4 = unit) & (e4 = op(e3, e4))) | ~ (e4 = op(e3, e4))) & ((~ (e4 = unit) & (e3 = op(e3, e4))) | ~ (e4 = op(e3, e3))) & ((~ (e4 = unit) & (e2 = op(e3, e4))) | ~ (e4 = op(e3, e2))) & ((~ (e4 = unit) & (e1 = op(e3, e4))) | ~ (e4 = op(e3, e1))) & ((~ (e4 = unit) & (e0 = op(e3, e4))) | ~ (e4 = op(e3, e0))) & ((~ (e4 = unit) & (e4 = op(e2, e4))) | ~ (e4 = op(e2, e4))) & ((~ (e4 = unit) & (e3 = op(e2, e4))) | ~ (e4 = op(e2, e3))) & ((~ (e4 = unit) & (e2 = op(e2, e4))) | ~ (e4 = op(e2, e2))) & ((~ (e4 = unit) & (e1 = op(e2, e4))) | ~ (e4 = op(e2, e1))) & ((~ (e4 = unit) & (e0 = op(e2, e4))) | ~ (e4 = op(e2, e0))) & ((~ (e4 = unit) & (e4 = op(e1, e4))) | ~ (e4 = op(e1, e4))) & ((~ (e4 = unit) & (e3 = op(e1, e4))) | ~ (e4 = op(e1, e3))) & ((~ (e4 = unit) & (e2 = op(e1, e4))) | ~ (e4 = op(e1, e2))) & ((~ (e4 = unit) & (e1 = op(e1, e4))) | ~ (e4 = op(e1, e1))) & ((~ (e4 = unit) & (e0 = op(e1, e4))) | ~ (e4 = op(e1, e0))) & ((~ (e4 = unit) & (e4 = op(e0, e4))) | ~ (e4 = op(e0, e4))) & ((~ (e4 = unit) & (e3 = op(e0, e4))) | ~ (e4 = op(e0, e3))) & ((~ (e4 = unit) & (e2 = op(e0, e4))) | ~ (e4 = op(e0, e2))) & ((~ (e4 = unit) & (e1 = op(e0, e4))) | ~ (e4 = op(e0, e1))) & ((~ (e4 = unit) & (e0 = op(e0, e4))) | ~ (e4 = op(e0, e0)))) | (((~ (e3 = unit) & (e4 = op(e4, e3))) | ~ (e3 = op(e4, e4))) & ((~ (e3 = unit) & (e3 = op(e4, e3))) | ~ (e3 = op(e4, e3))) & ((~ (e3 = unit) & (e2 = op(e4, e3))) | ~ (e3 = op(e4, e2))) & ((~ (e3 = unit) & (e1 = op(e4, e3))) | ~ (e3 = op(e4, e1))) & ((~ (e3 = unit) & (e0 = op(e4, e3))) | ~ (e3 = op(e4, e0))) & ((~ (e3 = unit) & (e4 = op(e3, e3))) | ~ (e3 = op(e3, e4))) & ((~ (e3 = unit) & (e3 = op(e3, e3))) | ~ (e3 = op(e3, e3))) & ((~ (e3 = unit) & (e2 = op(e3, e3))) | ~ (e3 = op(e3, e2))) & ((~ (e3 = unit) & (e1 = op(e3, e3))) | ~ (e3 = op(e3, e1))) & ((~ (e3 = unit) & (e0 = op(e3, e3))) | ~ (e3 = op(e3, e0))) & ((~ (e3 = unit) & (e4 = op(e2, e3))) | ~ (e3 = op(e2, e4))) & ((~ (e3 = unit) & (e3 = op(e2, e3))) | ~ (e3 = op(e2, e3))) & ((~ (e3 = unit) & (e2 = op(e2, e3))) | ~ (e3 = op(e2, e2))) & ((~ (e3 = unit) & (e1 = op(e2, e3))) | ~ (e3 = op(e2, e1))) & ((~ (e3 = unit) & (e0 = op(e2, e3))) | ~ (e3 = op(e2, e0))) & ((~ (e3 = unit) & (e4 = op(e1, e3))) | ~ (e3 = op(e1, e4))) & ((~ (e3 = unit) & (e3 = op(e1, e3))) | ~ (e3 = op(e1, e3))) & ((~ (e3 = unit) & (e2 = op(e1, e3))) | ~ (e3 = op(e1, e2))) & ((~ (e3 = unit) & (e1 = op(e1, e3))) | ~ (e3 = op(e1, e1))) & ((~ (e3 = unit) & (e0 = op(e1, e3))) | ~ (e3 = op(e1, e0))) & ((~ (e3 = unit) & (e4 = op(e0, e3))) | ~ (e3 = op(e0, e4))) & ((~ (e3 = unit) & (e3 = op(e0, e3))) | ~ (e3 = op(e0, e3))) & ((~ (e3 = unit) & (e2 = op(e0, e3))) | ~ (e3 = op(e0, e2))) & ((~ (e3 = unit) & (e1 = op(e0, e3))) | ~ (e3 = op(e0, e1))) & ((~ (e3 = unit) & (e0 = op(e0, e3))) | ~ (e3 = op(e0, e0)))) | (((~ (e2 = unit) & (e4 = op(e4, e2))) | ~ (e2 = op(e4, e4))) & ((~ (e2 = unit) & (e3 = op(e4, e2))) | ~ (e2 = op(e4, e3))) & ((~ (e2 = unit) & (e2 = op(e4, e2))) | ~ (e2 = op(e4, e2))) & ((~ (e2 = unit) & (e1 = op(e4, e2))) | ~ (e2 = op(e4, e1))) & ((~ (e2 = unit) & (e0 = op(e4, e2))) | ~ (e2 = op(e4, e0))) & ((~ (e2 = unit) & (e4 = op(e3, e2))) | ~ (e2 = op(e3, e4))) & ((~ (e2 = unit) & (e3 = op(e3, e2))) | ~ (e2 = op(e3, e3))) & ((~ (e2 = unit) & (e2 = op(e3, e2))) | ~ (e2 = op(e3, e2))) & ((~ (e2 = unit) & (e1 = op(e3, e2))) | ~ (e2 = op(e3, e1))) & ((~ (e2 = unit) & (e0 = op(e3, e2))) | ~ (e2 = op(e3, e0))) & ((~ (e2 = unit) & (e4 = op(e2, e2))) | ~ (e2 = op(e2, e4))) & ((~ (e2 = unit) & (e3 = op(e2, e2))) | ~ (e2 = op(e2, e3))) & ((~ (e2 = unit) & (e2 = op(e2, e2))) | ~ (e2 = op(e2, e2))) & ((~ (e2 = unit) & (e1 = op(e2, e2))) | ~ (e2 = op(e2, e1))) & ((~ (e2 = unit) & (e0 = op(e2, e2))) | ~ (e2 = op(e2, e0))) & ((~ (e2 = unit) & (e4 = op(e1, e2))) | ~ (e2 = op(e1, e4))) & ((~ (e2 = unit) & (e3 = op(e1, e2))) | ~ (e2 = op(e1, e3))) & ((~ (e2 = unit) & (e2 = op(e1, e2))) | ~ (e2 = op(e1, e2))) & ((~ (e2 = unit) & (e1 = op(e1, e2))) | ~ (e2 = op(e1, e1))) & ((~ (e2 = unit) & (e0 = op(e1, e2))) | ~ (e2 = op(e1, e0))) & ((~ (e2 = unit) & (e4 = op(e0, e2))) | ~ (e2 = op(e0, e4))) & ((~ (e2 = unit) & (e3 = op(e0, e2))) | ~ (e2 = op(e0, e3))) & ((~ (e2 = unit) & (e2 = op(e0, e2))) | ~ (e2 = op(e0, e2))) & ((~ (e2 = unit) & (e1 = op(e0, e2))) | ~ (e2 = op(e0, e1))) & ((~ (e2 = unit) & (e0 = op(e0, e2))) | ~ (e2 = op(e0, e0)))) | (((~ (e1 = unit) & (e4 = op(e4, e1))) | ~ (e1 = op(e4, e4))) & ((~ (e1 = unit) & (e3 = op(e4, e1))) | ~ (e1 = op(e4, e3))) & ((~ (e1 = unit) & (e2 = op(e4, e1))) | ~ (e1 = op(e4, e2))) & ((~ (e1 = unit) & (e1 = op(e4, e1))) | ~ (e1 = op(e4, e1))) & ((~ (e1 = unit) & (e0 = op(e4, e1))) | ~ (e1 = op(e4, e0))) & ((~ (e1 = unit) & (e4 = op(e3, e1))) | ~ (e1 = op(e3, e4))) & ((~ (e1 = unit) & (e3 = op(e3, e1))) | ~ (e1 = op(e3, e3))) & ((~ (e1 = unit) & (e2 = op(e3, e1))) | ~ (e1 = op(e3, e2))) & ((~ (e1 = unit) & (e1 = op(e3, e1))) | ~ (e1 = op(e3, e1))) & ((~ (e1 = unit) & (e0 = op(e3, e1))) | ~ (e1 = op(e3, e0))) & ((~ (e1 = unit) & (e4 = op(e2, e1))) | ~ (e1 = op(e2, e4))) & ((~ (e1 = unit) & (e3 = op(e2, e1))) | ~ (e1 = op(e2, e3))) & ((~ (e1 = unit) & (e2 = op(e2, e1))) | ~ (e1 = op(e2, e2))) & ((~ (e1 = unit) & (e1 = op(e2, e1))) | ~ (e1 = op(e2, e1))) & ((~ (e1 = unit) & (e0 = op(e2, e1))) | ~ (e1 = op(e2, e0))) & ((~ (e1 = unit) & (e4 = op(e1, e1))) | ~ (e1 = op(e1, e4))) & ((~ (e1 = unit) & (e3 = op(e1, e1))) | ~ (e1 = op(e1, e3))) & ((~ (e1 = unit) & (e2 = op(e1, e1))) | ~ (e1 = op(e1, e2))) & ((~ (e1 = unit) & (e1 = op(e1, e1))) | ~ (e1 = op(e1, e1))) & ((~ (e1 = unit) & (e0 = op(e1, e1))) | ~ (e1 = op(e1, e0))) & ((~ (e1 = unit) & (e4 = op(e0, e1))) | ~ (e1 = op(e0, e4))) & ((~ (e1 = unit) & (e3 = op(e0, e1))) | ~ (e1 = op(e0, e3))) & ((~ (e1 = unit) & (e2 = op(e0, e1))) | ~ (e1 = op(e0, e2))) & ((~ (e1 = unit) & (e1 = op(e0, e1))) | ~ (e1 = op(e0, e1))) & ((~ (e1 = unit) & (e0 = op(e0, e1))) | ~ (e1 = op(e0, e0)))) | (((~ (e0 = unit) & (e4 = op(e4, e0))) | ~ (e0 = op(e4, e4))) & ((~ (e0 = unit) & (e3 = op(e4, e0))) | ~ (e0 = op(e4, e3))) & ((~ (e0 = unit) & (e2 = op(e4, e0))) | ~ (e0 = op(e4, e2))) & ((~ (e0 = unit) & (e1 = op(e4, e0))) | ~ (e0 = op(e4, e1))) & ((~ (e0 = unit) & (e0 = op(e4, e0))) | ~ (e0 = op(e4, e0))) & ((~ (e0 = unit) & (e4 = op(e3, e0))) | ~ (e0 = op(e3, e4))) & ((~ (e0 = unit) & (e3 = op(e3, e0))) | ~ (e0 = op(e3, e3))) & ((~ (e0 = unit) & (e2 = op(e3, e0))) | ~ (e0 = op(e3, e2))) & ((~ (e0 = unit) & (e1 = op(e3, e0))) | ~ (e0 = op(e3, e1))) & ((~ (e0 = unit) & (e0 = op(e3, e0))) | ~ (e0 = op(e3, e0))) & ((~ (e0 = unit) & (e4 = op(e2, e0))) | ~ (e0 = op(e2, e4))) & ((~ (e0 = unit) & (e3 = op(e2, e0))) | ~ (e0 = op(e2, e3))) & ((~ (e0 = unit) & (e2 = op(e2, e0))) | ~ (e0 = op(e2, e2))) & ((~ (e0 = unit) & (e1 = op(e2, e0))) | ~ (e0 = op(e2, e1))) & ((~ (e0 = unit) & (e0 = op(e2, e0))) | ~ (e0 = op(e2, e0))) & ((~ (e0 = unit) & (e4 = op(e1, e0))) | ~ (e0 = op(e1, e4))) & ((~ (e0 = unit) & (e3 = op(e1, e0))) | ~ (e0 = op(e1, e3))) & ((~ (e0 = unit) & (e2 = op(e1, e0))) | ~ (e0 = op(e1, e2))) & ((~ (e0 = unit) & (e1 = op(e1, e0))) | ~ (e0 = op(e1, e1))) & ((~ (e0 = unit) & (e0 = op(e1, e0))) | ~ (e0 = op(e1, e0))) & ((~ (e0 = unit) & (e4 = op(e0, e0))) | ~ (e0 = op(e0, e4))) & ((~ (e0 = unit) & (e3 = op(e0, e0))) | ~ (e0 = op(e0, e3))) & ((~ (e0 = unit) & (e2 = op(e0, e0))) | ~ (e0 = op(e0, e2))) & ((~ (e0 = unit) & (e1 = op(e0, e0))) | ~ (e0 = op(e0, e1))) & ((~ (e0 = unit) & (e0 = op(e0, e0))) | ~ (e0 = op(e0, e0)))))), inference(negated_conjecture, [], [f4])).
fof(f4, plain, ~ (((e4 = op(e4, e4)) | (e4 = op(e3, e4)) | (e4 = op(e2, e4)) | (e4 = op(e1, e4)) | (e4 = op(e0, e4))) & ((e4 = op(e4, e4)) | (e4 = op(e4, e3)) | (e4 = op(e4, e2)) | (e4 = op(e4, e1)) | (e4 = op(e4, e0))) & ((e3 = op(e4, e4)) | (e3 = op(e3, e4)) | (e3 = op(e2, e4)) | (e3 = op(e1, e4)) | (e3 = op(e0, e4))) & ((e3 = op(e4, e4)) | (e3 = op(e4, e3)) | (e3 = op(e4, e2)) | (e3 = op(e4, e1)) | (e3 = op(e4, e0))) & ((e2 = op(e4, e4)) | (e2 = op(e3, e4)) | (e2 = op(e2, e4)) | (e2 = op(e1, e4)) | (e2 = op(e0, e4))) & ((e2 = op(e4, e4)) | (e2 = op(e4, e3)) | (e2 = op(e4, e2)) | (e2 = op(e4, e1)) | (e2 = op(e4, e0))) & ((e1 = op(e4, e4)) | (e1 = op(e3, e4)) | (e1 = op(e2, e4)) | (e1 = op(e1, e4)) | (e1 = op(e0, e4))) & ((e1 = op(e4, e4)) | (e1 = op(e4, e3)) | (e1 = op(e4, e2)) | (e1 = op(e4, e1)) | (e1 = op(e4, e0))) & ((e0 = op(e4, e4)) | (e0 = op(e3, e4)) | (e0 = op(e2, e4)) | (e0 = op(e1, e4)) | (e0 = op(e0, e4))) & ((e0 = op(e4, e4)) | (e0 = op(e4, e3)) | (e0 = op(e4, e2)) | (e0 = op(e4, e1)) | (e0 = op(e4, e0))) & ((e4 = op(e4, e3)) | (e4 = op(e3, e3)) | (e4 = op(e2, e3)) | (e4 = op(e1, e3)) | (e4 = op(e0, e3))) & ((e4 = op(e3, e4)) | (e4 = op(e3, e3)) | (e4 = op(e3, e2)) | (e4 = op(e3, e1)) | (e4 = op(e3, e0))) & ((e3 = op(e4, e3)) | (e3 = op(e3, e3)) | (e3 = op(e2, e3)) | (e3 = op(e1, e3)) | (e3 = op(e0, e3))) & ((e3 = op(e3, e4)) | (e3 = op(e3, e3)) | (e3 = op(e3, e2)) | (e3 = op(e3, e1)) | (e3 = op(e3, e0))) & ((e2 = op(e4, e3)) | (e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(e0, e3))) & ((e2 = op(e3, e4)) | (e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, e0))) & ((e1 = op(e4, e3)) | (e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(e0, e3))) & ((e1 = op(e3, e4)) | (e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, e0))) & ((e0 = op(e4, e3)) | (e0 = op(e3, e3)) | (e0 = op(e2, e3)) | (e0 = op(e1, e3)) | (e0 = op(e0, e3))) & ((e0 = op(e3, e4)) | (e0 = op(e3, e3)) | (e0 = op(e3, e2)) | (e0 = op(e3, e1)) | (e0 = op(e3, e0))) & ((e4 = op(e4, e2)) | (e4 = op(e3, e2)) | (e4 = op(e2, e2)) | (e4 = op(e1, e2)) | (e4 = op(e0, e2))) & ((e4 = op(e2, e4)) | (e4 = op(e2, e3)) | (e4 = op(e2, e2)) | (e4 = op(e2, e1)) | (e4 = op(e2, e0))) & ((e3 = op(e4, e2)) | (e3 = op(e3, e2)) | (e3 = op(e2, e2)) | (e3 = op(e1, e2)) | (e3 = op(e0, e2))) & ((e3 = op(e2, e4)) | (e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, e0))) & ((e2 = op(e4, e2)) | (e2 = op(e3, e2)) | (e2 = op(e2, e2)) | (e2 = op(e1, e2)) | (e2 = op(e0, e2))) & ((e2 = op(e2, e4)) | (e2 = op(e2, e3)) | (e2 = op(e2, e2)) | (e2 = op(e2, e1)) | (e2 = op(e2, e0))) & ((e1 = op(e4, e2)) | (e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(e0, e2))) & ((e1 = op(e2, e4)) | (e1 = op(e2, e3)) | (e1 = op(e2, e2)) | (e1 = op(e2, e1)) | (e1 = op(e2, e0))) & ((e0 = op(e4, e2)) | (e0 = op(e3, e2)) | (e0 = op(e2, e2)) | (e0 = op(e1, e2)) | (e0 = op(e0, e2))) & ((e0 = op(e2, e4)) | (e0 = op(e2, e3)) | (e0 = op(e2, e2)) | (e0 = op(e2, e1)) | (e0 = op(e2, e0))) & ((e4 = op(e4, e1)) | (e4 = op(e3, e1)) | (e4 = op(e2, e1)) | (e4 = op(e1, e1)) | (e4 = op(e0, e1))) & ((e4 = op(e1, e4)) | (e4 = op(e1, e3)) | (e4 = op(e1, e2)) | (e4 = op(e1, e1)) | (e4 = op(e1, e0))) & ((e3 = op(e4, e1)) | (e3 = op(e3, e1)) | (e3 = op(e2, e1)) | (e3 = op(e1, e1)) | (e3 = op(e0, e1))) & ((e3 = op(e1, e4)) | (e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, e0))) & ((e2 = op(e4, e1)) | (e2 = op(e3, e1)) | (e2 = op(e2, e1)) | (e2 = op(e1, e1)) | (e2 = op(e0, e1))) & ((e2 = op(e1, e4)) | (e2 = op(e1, e3)) | (e2 = op(e1, e2)) | (e2 = op(e1, e1)) | (e2 = op(e1, e0))) & ((e1 = op(e4, e1)) | (e1 = op(e3, e1)) | (e1 = op(e2, e1)) | (e1 = op(e1, e1)) | (e1 = op(e0, e1))) & ((e1 = op(e1, e4)) | (e1 = op(e1, e3)) | (e1 = op(e1, e2)) | (e1 = op(e1, e1)) | (e1 = op(e1, e0))) & ((e0 = op(e4, e1)) | (e0 = op(e3, e1)) | (e0 = op(e2, e1)) | (e0 = op(e1, e1)) | (e0 = op(e0, e1))) & ((e0 = op(e1, e4)) | (e0 = op(e1, e3)) | (e0 = op(e1, e2)) | (e0 = op(e1, e1)) | (e0 = op(e1, e0))) & ((e4 = op(e4, e0)) | (e4 = op(e3, e0)) | (e4 = op(e2, e0)) | (e4 = op(e1, e0)) | (e4 = op(e0, e0))) & ((e4 = op(e0, e4)) | (e4 = op(e0, e3)) | (e4 = op(e0, e2)) | (e4 = op(e0, e1)) | (e4 = op(e0, e0))) & ((e3 = op(e4, e0)) | (e3 = op(e3, e0)) | (e3 = op(e2, e0)) | (e3 = op(e1, e0)) | (e3 = op(e0, e0))) & ((e3 = op(e0, e4)) | (e3 = op(e0, e3)) | (e3 = op(e0, e2)) | (e3 = op(e0, e1)) | (e3 = op(e0, e0))) & ((e2 = op(e4, e0)) | (e2 = op(e3, e0)) | (e2 = op(e2, e0)) | (e2 = op(e1, e0)) | (e2 = op(e0, e0))) & ((e2 = op(e0, e4)) | (e2 = op(e0, e3)) | (e2 = op(e0, e2)) | (e2 = op(e0, e1)) | (e2 = op(e0, e0))) & ((e1 = op(e4, e0)) | (e1 = op(e3, e0)) | (e1 = op(e2, e0)) | (e1 = op(e1, e0)) | (e1 = op(e0, e0))) & ((e1 = op(e0, e4)) | (e1 = op(e0, e3)) | (e1 = op(e0, e2)) | (e1 = op(e0, e1)) | (e1 = op(e0, e0))) & ((e0 = op(e4, e0)) | (e0 = op(e3, e0)) | (e0 = op(e2, e0)) | (e0 = op(e1, e0)) | (e0 = op(e0, e0))) & ((e0 = op(e0, e4)) | (e0 = op(e0, e3)) | (e0 = op(e0, e2)) | (e0 = op(e0, e1)) | (e0 = op(e0, e0))) & ((e4 = unit) | (e3 = unit) | (e2 = unit) | (e1 = unit) | (e0 = unit)) & (e4 = op(e4, unit)) & (e4 = op(unit, e4)) & (e3 = op(e3, unit)) & (e3 = op(unit, e3)) & (e2 = op(e2, unit)) & (e2 = op(unit, e2)) & (e1 = op(e1, unit)) & (e1 = op(unit, e1)) & (e0 = op(e0, unit)) & (e0 = op(unit, e0)) & ((e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | (e1 = op(e4, e4)) | (e0 = op(e4, e4))) & ((e4 = op(e4, e3)) | (e3 = op(e4, e3)) | (e2 = op(e4, e3)) | (e1 = op(e4, e3)) | (e0 = op(e4, e3))) & ((e4 = op(e4, e2)) | (e3 = op(e4, e2)) | (e2 = op(e4, e2)) | (e1 = op(e4, e2)) | (e0 = op(e4, e2))) & ((e4 = op(e4, e1)) | (e3 = op(e4, e1)) | (e2 = op(e4, e1)) | (e1 = op(e4, e1)) | (e0 = op(e4, e1))) & ((e4 = op(e4, e0)) | (e3 = op(e4, e0)) | (e2 = op(e4, e0)) | (e1 = op(e4, e0)) | (e0 = op(e4, e0))) & ((e4 = op(e3, e4)) | (e3 = op(e3, e4)) | (e2 = op(e3, e4)) | (e1 = op(e3, e4)) | (e0 = op(e3, e4))) & ((e4 = op(e3, e3)) | (e3 = op(e3, e3)) | (e2 = op(e3, e3)) | (e1 = op(e3, e3)) | (e0 = op(e3, e3))) & ((e4 = op(e3, e2)) | (e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (e0 = op(e3, e2))) & ((e4 = op(e3, e1)) | (e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (e0 = op(e3, e1))) & ((e4 = op(e3, e0)) | (e3 = op(e3, e0)) | (e2 = op(e3, e0)) | (e1 = op(e3, e0)) | (e0 = op(e3, e0))) & ((e4 = op(e2, e4)) | (e3 = op(e2, e4)) | (e2 = op(e2, e4)) | (e1 = op(e2, e4)) | (e0 = op(e2, e4))) & ((e4 = op(e2, e3)) | (e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))) & ((e4 = op(e2, e2)) | (e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (e0 = op(e2, e2))) & ((e4 = op(e2, e1)) | (e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))) & ((e4 = op(e2, e0)) | (e3 = op(e2, e0)) | (e2 = op(e2, e0)) | (e1 = op(e2, e0)) | (e0 = op(e2, e0))) & ((e4 = op(e1, e4)) | (e3 = op(e1, e4)) | (e2 = op(e1, e4)) | (e1 = op(e1, e4)) | (e0 = op(e1, e4))) & ((e4 = op(e1, e3)) | (e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (e0 = op(e1, e3))) & ((e4 = op(e1, e2)) | (e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (e0 = op(e1, e2))) & ((e4 = op(e1, e1)) | (e3 = op(e1, e1)) | (e2 = op(e1, e1)) | (e1 = op(e1, e1)) | (e0 = op(e1, e1))) & ((e4 = op(e1, e0)) | (e3 = op(e1, e0)) | (e2 = op(e1, e0)) | (e1 = op(e1, e0)) | (e0 = op(e1, e0))) & ((e4 = op(e0, e4)) | (e3 = op(e0, e4)) | (e2 = op(e0, e4)) | (e1 = op(e0, e4)) | (e0 = op(e0, e4))) & ((e4 = op(e0, e3)) | (e3 = op(e0, e3)) | (e2 = op(e0, e3)) | (e1 = op(e0, e3)) | (e0 = op(e0, e3))) & ((e4 = op(e0, e2)) | (e3 = op(e0, e2)) | (e2 = op(e0, e2)) | (e1 = op(e0, e2)) | (e0 = op(e0, e2))) & ((e4 = op(e0, e1)) | (e3 = op(e0, e1)) | (e2 = op(e0, e1)) | (e1 = op(e0, e1)) | (e0 = op(e0, e1))) & ((e4 = op(e0, e0)) | (e3 = op(e0, e0)) | (e2 = op(e0, e0)) | (e1 = op(e0, e0)) | (e0 = op(e0, e0))) & ((((~ (e4 = unit) & (e4 = op(e4, e4))) | ~ (e4 = op(e4, e4))) & ((~ (e4 = unit) & (e3 = op(e4, e4))) | ~ (e4 = op(e4, e3))) & ((~ (e4 = unit) & (e2 = op(e4, e4))) | ~ (e4 = op(e4, e2))) & ((~ (e4 = unit) & (e1 = op(e4, e4))) | ~ (e4 = op(e4, e1))) & ((~ (e4 = unit) & (e0 = op(e4, e4))) | ~ (e4 = op(e4, e0))) & ((~ (e4 = unit) & (e4 = op(e3, e4))) | ~ (e4 = op(e3, e4))) & ((~ (e4 = unit) & (e3 = op(e3, e4))) | ~ (e4 = op(e3, e3))) & ((~ (e4 = unit) & (e2 = op(e3, e4))) | ~ (e4 = op(e3, e2))) & ((~ (e4 = unit) & (e1 = op(e3, e4))) | ~ (e4 = op(e3, e1))) & ((~ (e4 = unit) & (e0 = op(e3, e4))) | ~ (e4 = op(e3, e0))) & ((~ (e4 = unit) & (e4 = op(e2, e4))) | ~ (e4 = op(e2, e4))) & ((~ (e4 = unit) & (e3 = op(e2, e4))) | ~ (e4 = op(e2, e3))) & ((~ (e4 = unit) & (e2 = op(e2, e4))) | ~ (e4 = op(e2, e2))) & ((~ (e4 = unit) & (e1 = op(e2, e4))) | ~ (e4 = op(e2, e1))) & ((~ (e4 = unit) & (e0 = op(e2, e4))) | ~ (e4 = op(e2, e0))) & ((~ (e4 = unit) & (e4 = op(e1, e4))) | ~ (e4 = op(e1, e4))) & ((~ (e4 = unit) & (e3 = op(e1, e4))) | ~ (e4 = op(e1, e3))) & ((~ (e4 = unit) & (e2 = op(e1, e4))) | ~ (e4 = op(e1, e2))) & ((~ (e4 = unit) & (e1 = op(e1, e4))) | ~ (e4 = op(e1, e1))) & ((~ (e4 = unit) & (e0 = op(e1, e4))) | ~ (e4 = op(e1, e0))) & ((~ (e4 = unit) & (e4 = op(e0, e4))) | ~ (e4 = op(e0, e4))) & ((~ (e4 = unit) & (e3 = op(e0, e4))) | ~ (e4 = op(e0, e3))) & ((~ (e4 = unit) & (e2 = op(e0, e4))) | ~ (e4 = op(e0, e2))) & ((~ (e4 = unit) & (e1 = op(e0, e4))) | ~ (e4 = op(e0, e1))) & ((~ (e4 = unit) & (e0 = op(e0, e4))) | ~ (e4 = op(e0, e0)))) | (((~ (e3 = unit) & (e4 = op(e4, e3))) | ~ (e3 = op(e4, e4))) & ((~ (e3 = unit) & (e3 = op(e4, e3))) | ~ (e3 = op(e4, e3))) & ((~ (e3 = unit) & (e2 = op(e4, e3))) | ~ (e3 = op(e4, e2))) & ((~ (e3 = unit) & (e1 = op(e4, e3))) | ~ (e3 = op(e4, e1))) & ((~ (e3 = unit) & (e0 = op(e4, e3))) | ~ (e3 = op(e4, e0))) & ((~ (e3 = unit) & (e4 = op(e3, e3))) | ~ (e3 = op(e3, e4))) & ((~ (e3 = unit) & (e3 = op(e3, e3))) | ~ (e3 = op(e3, e3))) & ((~ (e3 = unit) & (e2 = op(e3, e3))) | ~ (e3 = op(e3, e2))) & ((~ (e3 = unit) & (e1 = op(e3, e3))) | ~ (e3 = op(e3, e1))) & ((~ (e3 = unit) & (e0 = op(e3, e3))) | ~ (e3 = op(e3, e0))) & ((~ (e3 = unit) & (e4 = op(e2, e3))) | ~ (e3 = op(e2, e4))) & ((~ (e3 = unit) & (e3 = op(e2, e3))) | ~ (e3 = op(e2, e3))) & ((~ (e3 = unit) & (e2 = op(e2, e3))) | ~ (e3 = op(e2, e2))) & ((~ (e3 = unit) & (e1 = op(e2, e3))) | ~ (e3 = op(e2, e1))) & ((~ (e3 = unit) & (e0 = op(e2, e3))) | ~ (e3 = op(e2, e0))) & ((~ (e3 = unit) & (e4 = op(e1, e3))) | ~ (e3 = op(e1, e4))) & ((~ (e3 = unit) & (e3 = op(e1, e3))) | ~ (e3 = op(e1, e3))) & ((~ (e3 = unit) & (e2 = op(e1, e3))) | ~ (e3 = op(e1, e2))) & ((~ (e3 = unit) & (e1 = op(e1, e3))) | ~ (e3 = op(e1, e1))) & ((~ (e3 = unit) & (e0 = op(e1, e3))) | ~ (e3 = op(e1, e0))) & ((~ (e3 = unit) & (e4 = op(e0, e3))) | ~ (e3 = op(e0, e4))) & ((~ (e3 = unit) & (e3 = op(e0, e3))) | ~ (e3 = op(e0, e3))) & ((~ (e3 = unit) & (e2 = op(e0, e3))) | ~ (e3 = op(e0, e2))) & ((~ (e3 = unit) & (e1 = op(e0, e3))) | ~ (e3 = op(e0, e1))) & ((~ (e3 = unit) & (e0 = op(e0, e3))) | ~ (e3 = op(e0, e0)))) | (((~ (e2 = unit) & (e4 = op(e4, e2))) | ~ (e2 = op(e4, e4))) & ((~ (e2 = unit) & (e3 = op(e4, e2))) | ~ (e2 = op(e4, e3))) & ((~ (e2 = unit) & (e2 = op(e4, e2))) | ~ (e2 = op(e4, e2))) & ((~ (e2 = unit) & (e1 = op(e4, e2))) | ~ (e2 = op(e4, e1))) & ((~ (e2 = unit) & (e0 = op(e4, e2))) | ~ (e2 = op(e4, e0))) & ((~ (e2 = unit) & (e4 = op(e3, e2))) | ~ (e2 = op(e3, e4))) & ((~ (e2 = unit) & (e3 = op(e3, e2))) | ~ (e2 = op(e3, e3))) & ((~ (e2 = unit) & (e2 = op(e3, e2))) | ~ (e2 = op(e3, e2))) & ((~ (e2 = unit) & (e1 = op(e3, e2))) | ~ (e2 = op(e3, e1))) & ((~ (e2 = unit) & (e0 = op(e3, e2))) | ~ (e2 = op(e3, e0))) & ((~ (e2 = unit) & (e4 = op(e2, e2))) | ~ (e2 = op(e2, e4))) & ((~ (e2 = unit) & (e3 = op(e2, e2))) | ~ (e2 = op(e2, e3))) & ((~ (e2 = unit) & (e2 = op(e2, e2))) | ~ (e2 = op(e2, e2))) & ((~ (e2 = unit) & (e1 = op(e2, e2))) | ~ (e2 = op(e2, e1))) & ((~ (e2 = unit) & (e0 = op(e2, e2))) | ~ (e2 = op(e2, e0))) & ((~ (e2 = unit) & (e4 = op(e1, e2))) | ~ (e2 = op(e1, e4))) & ((~ (e2 = unit) & (e3 = op(e1, e2))) | ~ (e2 = op(e1, e3))) & ((~ (e2 = unit) & (e2 = op(e1, e2))) | ~ (e2 = op(e1, e2))) & ((~ (e2 = unit) & (e1 = op(e1, e2))) | ~ (e2 = op(e1, e1))) & ((~ (e2 = unit) & (e0 = op(e1, e2))) | ~ (e2 = op(e1, e0))) & ((~ (e2 = unit) & (e4 = op(e0, e2))) | ~ (e2 = op(e0, e4))) & ((~ (e2 = unit) & (e3 = op(e0, e2))) | ~ (e2 = op(e0, e3))) & ((~ (e2 = unit) & (e2 = op(e0, e2))) | ~ (e2 = op(e0, e2))) & ((~ (e2 = unit) & (e1 = op(e0, e2))) | ~ (e2 = op(e0, e1))) & ((~ (e2 = unit) & (e0 = op(e0, e2))) | ~ (e2 = op(e0, e0)))) | (((~ (e1 = unit) & (e4 = op(e4, e1))) | ~ (e1 = op(e4, e4))) & ((~ (e1 = unit) & (e3 = op(e4, e1))) | ~ (e1 = op(e4, e3))) & ((~ (e1 = unit) & (e2 = op(e4, e1))) | ~ (e1 = op(e4, e2))) & ((~ (e1 = unit) & (e1 = op(e4, e1))) | ~ (e1 = op(e4, e1))) & ((~ (e1 = unit) & (e0 = op(e4, e1))) | ~ (e1 = op(e4, e0))) & ((~ (e1 = unit) & (e4 = op(e3, e1))) | ~ (e1 = op(e3, e4))) & ((~ (e1 = unit) & (e3 = op(e3, e1))) | ~ (e1 = op(e3, e3))) & ((~ (e1 = unit) & (e2 = op(e3, e1))) | ~ (e1 = op(e3, e2))) & ((~ (e1 = unit) & (e1 = op(e3, e1))) | ~ (e1 = op(e3, e1))) & ((~ (e1 = unit) & (e0 = op(e3, e1))) | ~ (e1 = op(e3, e0))) & ((~ (e1 = unit) & (e4 = op(e2, e1))) | ~ (e1 = op(e2, e4))) & ((~ (e1 = unit) & (e3 = op(e2, e1))) | ~ (e1 = op(e2, e3))) & ((~ (e1 = unit) & (e2 = op(e2, e1))) | ~ (e1 = op(e2, e2))) & ((~ (e1 = unit) & (e1 = op(e2, e1))) | ~ (e1 = op(e2, e1))) & ((~ (e1 = unit) & (e0 = op(e2, e1))) | ~ (e1 = op(e2, e0))) & ((~ (e1 = unit) & (e4 = op(e1, e1))) | ~ (e1 = op(e1, e4))) & ((~ (e1 = unit) & (e3 = op(e1, e1))) | ~ (e1 = op(e1, e3))) & ((~ (e1 = unit) & (e2 = op(e1, e1))) | ~ (e1 = op(e1, e2))) & ((~ (e1 = unit) & (e1 = op(e1, e1))) | ~ (e1 = op(e1, e1))) & ((~ (e1 = unit) & (e0 = op(e1, e1))) | ~ (e1 = op(e1, e0))) & ((~ (e1 = unit) & (e4 = op(e0, e1))) | ~ (e1 = op(e0, e4))) & ((~ (e1 = unit) & (e3 = op(e0, e1))) | ~ (e1 = op(e0, e3))) & ((~ (e1 = unit) & (e2 = op(e0, e1))) | ~ (e1 = op(e0, e2))) & ((~ (e1 = unit) & (e1 = op(e0, e1))) | ~ (e1 = op(e0, e1))) & ((~ (e1 = unit) & (e0 = op(e0, e1))) | ~ (e1 = op(e0, e0)))) | (((~ (e0 = unit) & (e4 = op(e4, e0))) | ~ (e0 = op(e4, e4))) & ((~ (e0 = unit) & (e3 = op(e4, e0))) | ~ (e0 = op(e4, e3))) & ((~ (e0 = unit) & (e2 = op(e4, e0))) | ~ (e0 = op(e4, e2))) & ((~ (e0 = unit) & (e1 = op(e4, e0))) | ~ (e0 = op(e4, e1))) & ((~ (e0 = unit) & (e0 = op(e4, e0))) | ~ (e0 = op(e4, e0))) & ((~ (e0 = unit) & (e4 = op(e3, e0))) | ~ (e0 = op(e3, e4))) & ((~ (e0 = unit) & (e3 = op(e3, e0))) | ~ (e0 = op(e3, e3))) & ((~ (e0 = unit) & (e2 = op(e3, e0))) | ~ (e0 = op(e3, e2))) & ((~ (e0 = unit) & (e1 = op(e3, e0))) | ~ (e0 = op(e3, e1))) & ((~ (e0 = unit) & (e0 = op(e3, e0))) | ~ (e0 = op(e3, e0))) & ((~ (e0 = unit) & (e4 = op(e2, e0))) | ~ (e0 = op(e2, e4))) & ((~ (e0 = unit) & (e3 = op(e2, e0))) | ~ (e0 = op(e2, e3))) & ((~ (e0 = unit) & (e2 = op(e2, e0))) | ~ (e0 = op(e2, e2))) & ((~ (e0 = unit) & (e1 = op(e2, e0))) | ~ (e0 = op(e2, e1))) & ((~ (e0 = unit) & (e0 = op(e2, e0))) | ~ (e0 = op(e2, e0))) & ((~ (e0 = unit) & (e4 = op(e1, e0))) | ~ (e0 = op(e1, e4))) & ((~ (e0 = unit) & (e3 = op(e1, e0))) | ~ (e0 = op(e1, e3))) & ((~ (e0 = unit) & (e2 = op(e1, e0))) | ~ (e0 = op(e1, e2))) & ((~ (e0 = unit) & (e1 = op(e1, e0))) | ~ (e0 = op(e1, e1))) & ((~ (e0 = unit) & (e0 = op(e1, e0))) | ~ (e0 = op(e1, e0))) & ((~ (e0 = unit) & (e4 = op(e0, e0))) | ~ (e0 = op(e0, e4))) & ((~ (e0 = unit) & (e3 = op(e0, e0))) | ~ (e0 = op(e0, e3))) & ((~ (e0 = unit) & (e2 = op(e0, e0))) | ~ (e0 = op(e0, e2))) & ((~ (e0 = unit) & (e1 = op(e0, e0))) | ~ (e0 = op(e0, e1))) & ((~ (e0 = unit) & (e0 = op(e0, e0))) | ~ (e0 = op(e0, e0)))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG093+1.p', co1)).
fof(f2401, plain, ~ spl191_51, inference(avatar_split_clause, [], [f2400, f1939])).
fof(f1939, plain, (spl191_51 <=> sP22), introduced(avatar_definition, [new_symbols(naming, [spl191_51])])).
fof(f2400, plain, ~ sP22, inference(subsumption_resolution, [], [f2399, f1074])).
fof(f1074, plain, ~ (e1 = unit), inference(definition_unfolding, [], [f390, f425])).
fof(f390, plain, ~ (e0 = e1), inference(cnf_transformation, [], [f1])).
fof(f1, plain, (~ (e3 = e4) & ~ (e2 = e4) & ~ (e2 = e3) & ~ (e1 = e4) & ~ (e1 = e3) & ~ (e1 = e2) & ~ (e0 = e4) & ~ (e0 = e3) & ~ (e0 = e2) & ~ (e0 = e1)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG093+1.p', ax1)).
fof(f2399, plain, ((e1 = unit) | ~ sP22), inference(forward_demodulation, [], [f1350, f1087])).
fof(f1350, plain, ((e1 = op(unit, unit)) | ~ sP22), inference(definition_unfolding, [], [f1020, f425, f425])).
fof(f1020, plain, ((e1 = op(e0, e0)) | ~ sP22), inference(cnf_transformation, [], [f367])).
fof(f367, plain, ((((e1 = unit) | ~ (e0 = op(e0, e1))) & (e1 = op(e0, e0))) | ~ sP22), inference(nnf_transformation, [], [f29])).
fof(f2396, plain, ~ spl191_52, inference(avatar_split_clause, [], [f2395, f1943])).
fof(f1943, plain, (spl191_52 <=> sP23), introduced(avatar_definition, [new_symbols(naming, [spl191_52])])).
fof(f2395, plain, ~ sP23, inference(subsumption_resolution, [], [f2394, f1086])).
fof(f2394, plain, (~ (e1 = op(unit, e1)) | ~ sP23), inference(subsumption_resolution, [], [f1347, f1074])).
fof(f1347, plain, ((e1 = unit) | ~ (e1 = op(unit, e1)) | ~ sP23), inference(definition_unfolding, [], [f1019, f425])).
fof(f1019, plain, ((e1 = unit) | ~ (e1 = op(e0, e1)) | ~ sP23), inference(cnf_transformation, [], [f366])).
fof(f366, plain, ((((e1 = unit) | ~ (e1 = op(e0, e1))) & (e1 = op(e0, e1))) | ~ sP23), inference(nnf_transformation, [], [f30])).
fof(f2393, plain, ~ spl191_53, inference(avatar_split_clause, [], [f2392, f1947])).
fof(f1947, plain, (spl191_53 <=> sP24), introduced(avatar_definition, [new_symbols(naming, [spl191_53])])).
fof(f2392, plain, ~ sP24, inference(subsumption_resolution, [], [f2391, f394])).
fof(f394, plain, ~ (e1 = e2), inference(cnf_transformation, [], [f1])).
fof(f2391, plain, ((e1 = e2) | ~ sP24), inference(forward_demodulation, [], [f1346, f1085])).
fof(f1346, plain, ((e1 = op(unit, e2)) | ~ sP24), inference(definition_unfolding, [], [f1016, f425])).
fof(f1016, plain, ((e1 = op(e0, e2)) | ~ sP24), inference(cnf_transformation, [], [f365])).
fof(f365, plain, ((((e1 = unit) | ~ (e2 = op(e0, e1))) & (e1 = op(e0, e2))) | ~ sP24), inference(nnf_transformation, [], [f31])).
fof(f2388, plain, ~ spl191_54, inference(avatar_split_clause, [], [f2387, f1951])).
fof(f1951, plain, (spl191_54 <=> sP25), introduced(avatar_definition, [new_symbols(naming, [spl191_54])])).
fof(f2387, plain, ~ sP25, inference(subsumption_resolution, [], [f2386, f395])).
fof(f395, plain, ~ (e1 = e3), inference(cnf_transformation, [], [f1])).
fof(f2386, plain, ((e1 = e3) | ~ sP25), inference(forward_demodulation, [], [f1344, f1084])).
fof(f1344, plain, ((e1 = op(unit, e3)) | ~ sP25), inference(definition_unfolding, [], [f1014, f425])).
fof(f1014, plain, ((e1 = op(e0, e3)) | ~ sP25), inference(cnf_transformation, [], [f364])).
fof(f364, plain, ((((e1 = unit) | ~ (e3 = op(e0, e1))) & (e1 = op(e0, e3))) | ~ sP25), inference(nnf_transformation, [], [f32])).
fof(f2383, plain, ~ spl191_55, inference(avatar_split_clause, [], [f2382, f1955])).
fof(f1955, plain, (spl191_55 <=> sP26), introduced(avatar_definition, [new_symbols(naming, [spl191_55])])).
fof(f2382, plain, ~ sP26, inference(subsumption_resolution, [], [f2381, f396])).
fof(f396, plain, ~ (e1 = e4), inference(cnf_transformation, [], [f1])).
fof(f2381, plain, ((e1 = e4) | ~ sP26), inference(forward_demodulation, [], [f1342, f1083])).
fof(f1342, plain, ((e1 = op(unit, e4)) | ~ sP26), inference(definition_unfolding, [], [f1012, f425])).
fof(f1012, plain, ((e1 = op(e0, e4)) | ~ sP26), inference(cnf_transformation, [], [f363])).
fof(f363, plain, ((((e1 = unit) | ~ (e4 = op(e0, e1))) & (e1 = op(e0, e4))) | ~ sP26), inference(nnf_transformation, [], [f33])).
fof(f2378, plain, ~ spl191_56, inference(avatar_split_clause, [], [f2377, f1959])).
fof(f1959, plain, (spl191_56 <=> sP27), introduced(avatar_definition, [new_symbols(naming, [spl191_56])])).
fof(f2377, plain, ~ sP27, inference(subsumption_resolution, [], [f2376, f1081])).
fof(f2376, plain, (~ (op(e1, e1) = unit) | ~ sP27), inference(subsumption_resolution, [], [f1339, f1074])).
fof(f1339, plain, ((e1 = unit) | ~ (op(e1, e1) = unit) | ~ sP27), inference(definition_unfolding, [], [f1011, f425])).
fof(f1011, plain, ((e1 = unit) | ~ (e0 = op(e1, e1)) | ~ sP27), inference(cnf_transformation, [], [f362])).
fof(f362, plain, ((((e1 = unit) | ~ (e0 = op(e1, e1))) & (e1 = op(e1, e0))) | ~ sP27), inference(nnf_transformation, [], [f34])).
fof(f2375, plain, ~ spl191_57, inference(avatar_split_clause, [], [f2374, f1963])).
fof(f1963, plain, (spl191_57 <=> sP28), introduced(avatar_definition, [new_symbols(naming, [spl191_57])])).
fof(f2374, plain, ~ sP28, inference(subsumption_resolution, [], [f2373, f1074])).
fof(f2373, plain, ((e1 = unit) | ~ sP28), inference(forward_demodulation, [], [f1008, f1081])).
fof(f1008, plain, ((e1 = op(e1, e1)) | ~ sP28), inference(cnf_transformation, [], [f361])).
fof(f361, plain, ((((e1 = unit) | ~ (e1 = op(e1, e1))) & (e1 = op(e1, e1))) | ~ sP28), inference(nnf_transformation, [], [f35])).
fof(f2370, plain, ~ spl191_58, inference(avatar_split_clause, [], [f2369, f1967])).
fof(f1967, plain, (spl191_58 <=> sP29), introduced(avatar_definition, [new_symbols(naming, [spl191_58])])).
fof(f2369, plain, ~ sP29, inference(subsumption_resolution, [], [f2368, f396])).
fof(f2368, plain, ((e1 = e4) | ~ sP29), inference(forward_demodulation, [], [f1006, f407])).
fof(f1006, plain, ((e1 = op(e1, e2)) | ~ sP29), inference(cnf_transformation, [], [f360])).
fof(f360, plain, ((((e1 = unit) | ~ (e2 = op(e1, e1))) & (e1 = op(e1, e2))) | ~ sP29), inference(nnf_transformation, [], [f36])).
fof(f2365, plain, ~ spl191_59, inference(avatar_split_clause, [], [f2364, f1971])).
fof(f1971, plain, (spl191_59 <=> sP30), introduced(avatar_definition, [new_symbols(naming, [spl191_59])])).
fof(f2364, plain, ~ sP30, inference(subsumption_resolution, [], [f2363, f394])).
fof(f2363, plain, ((e1 = e2) | ~ sP30), inference(forward_demodulation, [], [f1004, f408])).
fof(f1004, plain, ((e1 = op(e1, e3)) | ~ sP30), inference(cnf_transformation, [], [f359])).
fof(f359, plain, ((((e1 = unit) | ~ (e3 = op(e1, e1))) & (e1 = op(e1, e3))) | ~ sP30), inference(nnf_transformation, [], [f37])).
fof(f2360, plain, ~ spl191_60, inference(avatar_split_clause, [], [f2359, f1975])).
fof(f1975, plain, (spl191_60 <=> sP31), introduced(avatar_definition, [new_symbols(naming, [spl191_60])])).
fof(f2359, plain, ~ sP31, inference(subsumption_resolution, [], [f2358, f395])).
fof(f2358, plain, ((e1 = e3) | ~ sP31), inference(forward_demodulation, [], [f1002, f409])).
fof(f1002, plain, ((e1 = op(e1, e4)) | ~ sP31), inference(cnf_transformation, [], [f358])).
fof(f358, plain, ((((e1 = unit) | ~ (e4 = op(e1, e1))) & (e1 = op(e1, e4))) | ~ sP31), inference(nnf_transformation, [], [f38])).
fof(f2355, plain, ~ spl191_61, inference(avatar_split_clause, [], [f2354, f1979])).
fof(f1979, plain, (spl191_61 <=> sP32), introduced(avatar_definition, [new_symbols(naming, [spl191_61])])).
fof(f2354, plain, ~ sP32, inference(subsumption_resolution, [], [f2353, f394])).
fof(f2353, plain, ((e1 = e2) | ~ sP32), inference(forward_demodulation, [], [f1338, f1080])).
fof(f1338, plain, ((e1 = op(e2, unit)) | ~ sP32), inference(definition_unfolding, [], [f1000, f425])).
fof(f1000, plain, ((e1 = op(e2, e0)) | ~ sP32), inference(cnf_transformation, [], [f357])).
fof(f357, plain, ((((e1 = unit) | ~ (e0 = op(e2, e1))) & (e1 = op(e2, e0))) | ~ sP32), inference(nnf_transformation, [], [f39])).
fof(f2350, plain, ~ spl191_62, inference(avatar_split_clause, [], [f2349, f1983])).
fof(f1983, plain, (spl191_62 <=> sP33), introduced(avatar_definition, [new_symbols(naming, [spl191_62])])).
fof(f2349, plain, ~ sP33, inference(subsumption_resolution, [], [f2348, f396])).
fof(f2348, plain, ((e1 = e4) | ~ sP33), inference(forward_demodulation, [], [f998, f411])).
fof(f998, plain, ((e1 = op(e2, e1)) | ~ sP33), inference(cnf_transformation, [], [f356])).
fof(f356, plain, ((((e1 = unit) | ~ (e1 = op(e2, e1))) & (e1 = op(e2, e1))) | ~ sP33), inference(nnf_transformation, [], [f40])).
fof(f2345, plain, ~ spl191_63, inference(avatar_split_clause, [], [f2344, f1987])).
fof(f1987, plain, (spl191_63 <=> sP34), introduced(avatar_definition, [new_symbols(naming, [spl191_63])])).
fof(f2344, plain, ~ sP34, inference(subsumption_resolution, [], [f2343, f395])).
fof(f2343, plain, ((e1 = e3) | ~ sP34), inference(forward_demodulation, [], [f996, f412])).
fof(f996, plain, ((e1 = op(e2, e2)) | ~ sP34), inference(cnf_transformation, [], [f355])).
fof(f355, plain, ((((e1 = unit) | ~ (e2 = op(e2, e1))) & (e1 = op(e2, e2))) | ~ sP34), inference(nnf_transformation, [], [f41])).
fof(f2340, plain, ~ spl191_64, inference(avatar_split_clause, [], [f2339, f1991])).
fof(f1991, plain, (spl191_64 <=> sP35), introduced(avatar_definition, [new_symbols(naming, [spl191_64])])).
fof(f2339, plain, ~ sP35, inference(subsumption_resolution, [], [f2338, f1074])).
fof(f2338, plain, ((e1 = unit) | ~ sP35), inference(forward_demodulation, [], [f994, f1079])).
fof(f994, plain, ((e1 = op(e2, e3)) | ~ sP35), inference(cnf_transformation, [], [f354])).
fof(f354, plain, ((((e1 = unit) | ~ (e3 = op(e2, e1))) & (e1 = op(e2, e3))) | ~ sP35), inference(nnf_transformation, [], [f42])).
fof(f2335, plain, ~ spl191_65, inference(avatar_split_clause, [], [f2334, f1995])).
fof(f1995, plain, (spl191_65 <=> sP36), introduced(avatar_definition, [new_symbols(naming, [spl191_65])])).
fof(f2334, plain, ~ sP36, inference(subsumption_resolution, [], [f2333, f411])).
fof(f2333, plain, (~ (e4 = op(e2, e1)) | ~ sP36), inference(subsumption_resolution, [], [f993, f1074])).
fof(f993, plain, ((e1 = unit) | ~ (e4 = op(e2, e1)) | ~ sP36), inference(cnf_transformation, [], [f353])).
fof(f353, plain, ((((e1 = unit) | ~ (e4 = op(e2, e1))) & (e1 = op(e2, e4))) | ~ sP36), inference(nnf_transformation, [], [f43])).
fof(f2332, plain, ~ spl191_66, inference(avatar_split_clause, [], [f2331, f1999])).
fof(f1999, plain, (spl191_66 <=> sP37), introduced(avatar_definition, [new_symbols(naming, [spl191_66])])).
fof(f2331, plain, ~ sP37, inference(subsumption_resolution, [], [f2330, f395])).
fof(f2330, plain, ((e1 = e3) | ~ sP37), inference(forward_demodulation, [], [f1336, f1078])).
fof(f1336, plain, ((e1 = op(e3, unit)) | ~ sP37), inference(definition_unfolding, [], [f990, f425])).
fof(f990, plain, ((e1 = op(e3, e0)) | ~ sP37), inference(cnf_transformation, [], [f352])).
fof(f352, plain, ((((e1 = unit) | ~ (e0 = op(e3, e1))) & (e1 = op(e3, e0))) | ~ sP37), inference(nnf_transformation, [], [f44])).
fof(f2327, plain, ~ spl191_67, inference(avatar_split_clause, [], [f2326, f2003])).
fof(f2003, plain, (spl191_67 <=> sP38), introduced(avatar_definition, [new_symbols(naming, [spl191_67])])).
fof(f2326, plain, ~ sP38, inference(subsumption_resolution, [], [f2325, f394])).
fof(f2325, plain, ((e1 = e2) | ~ sP38), inference(forward_demodulation, [], [f988, f416])).
fof(f988, plain, ((e1 = op(e3, e1)) | ~ sP38), inference(cnf_transformation, [], [f351])).
fof(f351, plain, ((((e1 = unit) | ~ (e1 = op(e3, e1))) & (e1 = op(e3, e1))) | ~ sP38), inference(nnf_transformation, [], [f45])).
fof(f2322, plain, ~ spl191_68, inference(avatar_split_clause, [], [f2321, f2007])).
fof(f2007, plain, (spl191_68 <=> sP39), introduced(avatar_definition, [new_symbols(naming, [spl191_68])])).
fof(f2321, plain, ~ sP39, inference(subsumption_resolution, [], [f2320, f416])).
fof(f2320, plain, (~ (e2 = op(e3, e1)) | ~ sP39), inference(subsumption_resolution, [], [f987, f1074])).
fof(f987, plain, ((e1 = unit) | ~ (e2 = op(e3, e1)) | ~ sP39), inference(cnf_transformation, [], [f350])).
fof(f350, plain, ((((e1 = unit) | ~ (e2 = op(e3, e1))) & (e1 = op(e3, e2))) | ~ sP39), inference(nnf_transformation, [], [f46])).
fof(f2319, plain, ~ spl191_69, inference(avatar_split_clause, [], [f2318, f2011])).
fof(f2011, plain, (spl191_69 <=> sP40), introduced(avatar_definition, [new_symbols(naming, [spl191_69])])).
fof(f2318, plain, ~ sP40, inference(subsumption_resolution, [], [f2317, f396])).
fof(f2317, plain, ((e1 = e4) | ~ sP40), inference(forward_demodulation, [], [f984, f418])).
fof(f984, plain, ((e1 = op(e3, e3)) | ~ sP40), inference(cnf_transformation, [], [f349])).
fof(f349, plain, ((((e1 = unit) | ~ (e3 = op(e3, e1))) & (e1 = op(e3, e3))) | ~ sP40), inference(nnf_transformation, [], [f47])).
fof(f2314, plain, ~ spl191_70, inference(avatar_split_clause, [], [f2313, f2015])).
fof(f2015, plain, (spl191_70 <=> sP41), introduced(avatar_definition, [new_symbols(naming, [spl191_70])])).
fof(f2313, plain, ~ sP41, inference(subsumption_resolution, [], [f2312, f1074])).
fof(f2312, plain, ((e1 = unit) | ~ sP41), inference(forward_demodulation, [], [f982, f1077])).
fof(f982, plain, ((e1 = op(e3, e4)) | ~ sP41), inference(cnf_transformation, [], [f348])).
fof(f348, plain, ((((e1 = unit) | ~ (e4 = op(e3, e1))) & (e1 = op(e3, e4))) | ~ sP41), inference(nnf_transformation, [], [f48])).
fof(f2309, plain, ~ spl191_71, inference(avatar_split_clause, [], [f2308, f2019])).
fof(f2019, plain, (spl191_71 <=> sP42), introduced(avatar_definition, [new_symbols(naming, [spl191_71])])).
fof(f2308, plain, ~ sP42, inference(subsumption_resolution, [], [f2307, f396])).
fof(f2307, plain, ((e1 = e4) | ~ sP42), inference(forward_demodulation, [], [f1334, f1076])).
fof(f1334, plain, ((e1 = op(e4, unit)) | ~ sP42), inference(definition_unfolding, [], [f980, f425])).
fof(f980, plain, ((e1 = op(e4, e0)) | ~ sP42), inference(cnf_transformation, [], [f347])).
fof(f347, plain, ((((e1 = unit) | ~ (e0 = op(e4, e1))) & (e1 = op(e4, e0))) | ~ sP42), inference(nnf_transformation, [], [f49])).
fof(f2304, plain, ~ spl191_72, inference(avatar_split_clause, [], [f2303, f2023])).
fof(f2023, plain, (spl191_72 <=> sP43), introduced(avatar_definition, [new_symbols(naming, [spl191_72])])).
fof(f2303, plain, ~ sP43, inference(subsumption_resolution, [], [f2302, f395])).
fof(f2302, plain, ((e1 = e3) | ~ sP43), inference(forward_demodulation, [], [f978, f421])).
fof(f978, plain, ((e1 = op(e4, e1)) | ~ sP43), inference(cnf_transformation, [], [f346])).
fof(f346, plain, ((((e1 = unit) | ~ (e1 = op(e4, e1))) & (e1 = op(e4, e1))) | ~ sP43), inference(nnf_transformation, [], [f50])).
fof(f2026, plain, (~ spl191_5 | spl191_51 | spl191_52 | spl191_53 | spl191_54 | spl191_55 | spl191_56 | spl191_57 | spl191_58 | spl191_59 | spl191_60 | spl191_61 | spl191_62 | spl191_63 | spl191_64 | spl191_65 | spl191_66 | spl191_67 | spl191_68 | spl191_69 | spl191_70 | spl191_71 | spl191_72), inference(avatar_split_clause, [], [f1937, f2023, f2019, f2015, f2011, f2007, f2003, f1999, f1995, f1991, f1987, f1983, f1979, f1975, f1971, f1967, f1963, f1959, f1955, f1951, f1947, f1943, f1939, f1666])).
fof(f1666, plain, (spl191_5 <=> sP111), introduced(avatar_definition, [new_symbols(naming, [spl191_5])])).
fof(f1937, plain, (sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | ~ sP111), inference(subsumption_resolution, [], [f1936, f1074])).
fof(f1936, plain, ((e1 = unit) | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | ~ sP111), inference(forward_demodulation, [], [f1935, f1075])).
fof(f1935, plain, ((e1 = op(e4, e2)) | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | ~ sP111), inference(subsumption_resolution, [], [f1934, f394])).
fof(f1934, plain, ((e1 = e2) | (e1 = op(e4, e2)) | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | ~ sP111), inference(forward_demodulation, [], [f1933, f424])).
fof(f1933, plain, ((e1 = op(e4, e4)) | (e1 = op(e4, e2)) | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | ~ sP111), inference(subsumption_resolution, [], [f1932, f421])).
fof(f1932, plain, ((e1 = op(e4, e4)) | ~ (e3 = op(e4, e1)) | (e1 = op(e4, e2)) | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | ~ sP111), inference(subsumption_resolution, [], [f832, f1074])).
fof(f832, plain, ((e1 = op(e4, e4)) | (e1 = unit) | ~ (e3 = op(e4, e1)) | (e1 = op(e4, e2)) | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | ~ sP111), inference(cnf_transformation, [], [f278])).
fof(f278, plain, ((((e1 = unit) | ~ (e4 = op(e4, e1))) & (e1 = op(e4, e4))) | (((e1 = unit) | ~ (e3 = op(e4, e1))) & (e1 = op(e4, e3))) | (((e1 = unit) | ~ (e2 = op(e4, e1))) & (e1 = op(e4, e2))) | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | ~ sP111), inference(nnf_transformation, [], [f118])).
fof(f1669, plain, (~ spl191_1 | spl191_5), inference(avatar_split_clause, [], [f802, f1666, f1647])).
fof(f802, plain, (sP111 | ~ sP115), inference(cnf_transformation, [], [f274])).
fof(f274, plain, ((sP114 & sP113 & sP112 & sP111 & sP110) | ~ sP115), inference(nnf_transformation, [], [f122])).