fof(f3323, plain, $false, inference(avatar_sat_refutation, [], [f374, f408, f425, f544, f578, f663, f664, f665, f666, f673, f674, f679, f680, f682, f683, f684, f685, f686, f687, f698, f703, f706, f848, f849, f851, f852, f853, f940, f1146, f1148, f1151, f1176, f1253, f1254, f1614, f1663, f1795, f1938, f1972, f2013, f2027, f2076, f2147, f2221, f2260, f2291, f2556, f2568, f2598, f2634, f2744, f2747, f2751, f2752, f2787, f2956, f2989, f3107, f3108, f3133, f3203, f3204, f3225])).
fof(f3225, plain, (~ spl3_17 | ~ spl3_21), inference(avatar_split_clause, [], [f3211, f376, f359])).
fof(f359, plain, (spl3_17 <=> (e0 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl3_17])])).
fof(f376, plain, (spl3_21 <=> (e0 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl3_21])])).
fof(f3211, plain, (~ (e0 = op(e2, e3)) | ~ spl3_21), inference(backward_demodulation, [], [f237, f378])).
fof(f378, plain, ((e0 = op(e2, e2)) | ~ spl3_21), inference(avatar_component_clause, [], [f376])).
fof(f237, plain, ~ (op(e2, e2) = op(e2, e3)), inference(cnf_transformation, [], [f9])).
fof(f9, plain, (~ (op(e3, e2) = op(e3, e3)) & ~ (op(e3, e1) = op(e3, e3)) & ~ (op(e3, e1) = op(e3, e2)) & ~ (op(e3, e0) = op(e3, e3)) & ~ (op(e3, e0) = op(e3, e2)) & ~ (op(e3, e0) = op(e3, e1)) & ~ (op(e2, e2) = op(e2, e3)) & ~ (op(e2, e1) = op(e2, e3)) & ~ (op(e2, e1) = op(e2, e2)) & ~ (op(e2, e0) = op(e2, e3)) & ~ (op(e2, e0) = op(e2, e2)) & ~ (op(e2, e0) = op(e2, e1)) & ~ (op(e1, e2) = op(e1, e3)) & ~ (op(e1, e1) = op(e1, e3)) & ~ (op(e1, e1) = op(e1, e2)) & ~ (op(e1, e0) = op(e1, e3)) & ~ (op(e1, e0) = op(e1, e2)) & ~ (op(e1, e0) = op(e1, e1)) & ~ (op(e0, e2) = op(e0, e3)) & ~ (op(e0, e1) = op(e0, e3)) & ~ (op(e0, e1) = op(e0, e2)) & ~ (op(e0, e0) = op(e0, e3)) & ~ (op(e0, e0) = op(e0, e2)) & ~ (op(e0, e0) = op(e0, e1)) & ~ (op(e2, e3) = op(e3, e3)) & ~ (op(e1, e3) = op(e3, e3)) & ~ (op(e1, e3) = op(e2, e3)) & ~ (op(e0, e3) = op(e3, e3)) & ~ (op(e0, e3) = op(e2, e3)) & ~ (op(e0, e3) = op(e1, e3)) & ~ (op(e2, e2) = op(e3, e2)) & ~ (op(e1, e2) = op(e3, e2)) & ~ (op(e1, e2) = op(e2, e2)) & ~ (op(e0, e2) = op(e3, e2)) & ~ (op(e0, e2) = op(e2, e2)) & ~ (op(e0, e2) = op(e1, e2)) & ~ (op(e2, e1) = op(e3, e1)) & ~ (op(e1, e1) = op(e3, e1)) & ~ (op(e1, e1) = op(e2, e1)) & ~ (op(e0, e1) = op(e3, e1)) & ~ (op(e0, e1) = op(e2, e1)) & ~ (op(e0, e1) = op(e1, e1)) & ~ (op(e2, e0) = op(e3, e0)) & ~ (op(e1, e0) = op(e3, e0)) & ~ (op(e1, e0) = op(e2, e0)) & ~ (op(e0, e0) = op(e3, e0)) & ~ (op(e0, e0) = op(e2, e0)) & ~ (op(e0, e0) = op(e1, e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG023+1.p', ax9)).
fof(f3204, plain, (~ spl3_19 | ~ spl3_27), inference(avatar_split_clause, [], [f3189, f401, f367])).
fof(f367, plain, (spl3_19 <=> (e2 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl3_19])])).
fof(f401, plain, (spl3_27 <=> (e2 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl3_27])])).
fof(f3189, plain, (~ (e2 = op(e2, e3)) | ~ spl3_27), inference(backward_demodulation, [], [f236, f403])).
fof(f403, plain, ((e2 = op(e2, e1)) | ~ spl3_27), inference(avatar_component_clause, [], [f401])).
fof(f236, plain, ~ (op(e2, e1) = op(e2, e3)), inference(cnf_transformation, [], [f9])).
fof(f3203, plain, (~ spl3_23 | ~ spl3_27), inference(avatar_split_clause, [], [f3188, f401, f384])).
fof(f384, plain, (spl3_23 <=> (e2 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl3_23])])).
fof(f3188, plain, (~ (e2 = op(e2, e2)) | ~ spl3_27), inference(backward_demodulation, [], [f235, f403])).
fof(f235, plain, ~ (op(e2, e1) = op(e2, e2)), inference(cnf_transformation, [], [f9])).
fof(f3133, plain, (~ spl3_13 | ~ spl3_61), inference(avatar_split_clause, [], [f3125, f546, f342])).
fof(f342, plain, (spl3_13 <=> (e0 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl3_13])])).
fof(f546, plain, (spl3_61 <=> (e0 = op(e0, e0))), introduced(avatar_definition, [new_symbols(naming, [spl3_61])])).
fof(f3125, plain, (~ (e0 = op(e3, e0)) | ~ spl3_61), inference(backward_demodulation, [], [f198, f548])).
fof(f548, plain, ((e0 = op(e0, e0)) | ~ spl3_61), inference(avatar_component_clause, [], [f546])).
fof(f198, plain, ~ (op(e0, e0) = op(e3, e0)), inference(cnf_transformation, [], [f9])).
fof(f3108, plain, (spl3_27 | ~ spl3_66), inference(avatar_split_clause, [], [f3092, f567, f401])).
fof(f567, plain, (spl3_66 <=> (e1 = unit)), introduced(avatar_definition, [new_symbols(naming, [spl3_66])])).
fof(f3092, plain, ((e2 = op(e2, e1)) | ~ spl3_66), inference(backward_demodulation, [], [f153, f569])).
fof(f569, plain, ((e1 = unit) | ~ spl3_66), inference(avatar_component_clause, [], [f567])).
fof(f153, plain, (e2 = op(e2, unit)), inference(cnf_transformation, [], [f3])).
fof(f3, plain, (((e3 = unit) | (e2 = unit) | (e1 = unit) | (e0 = unit)) & (e3 = op(e3, unit)) & (e3 = op(unit, e3)) & (e2 = op(e2, unit)) & (e2 = op(unit, e2)) & (e1 = op(e1, unit)) & (e1 = op(unit, e1)) & (e0 = op(e0, unit)) & (e0 = op(unit, e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG023+1.p', ax3)).
fof(f3107, plain, (spl3_45 | ~ spl3_66), inference(avatar_split_clause, [], [f3088, f567, f478])).
fof(f478, plain, (spl3_45 <=> (e0 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl3_45])])).
fof(f3088, plain, ((e0 = op(e1, e0)) | ~ spl3_66), inference(backward_demodulation, [], [f148, f569])).
fof(f148, plain, (e0 = op(unit, e0)), inference(cnf_transformation, [], [f3])).
fof(f2989, plain, (~ spl3_61 | ~ spl3_57), inference(avatar_split_clause, [], [f2988, f529, f546])).
fof(f529, plain, (spl3_57 <=> (e0 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl3_57])])).
fof(f2988, plain, (~ (e0 = op(e0, e0)) | ~ spl3_57), inference(forward_demodulation, [], [f220, f531])).
fof(f531, plain, ((e0 = op(e0, e1)) | ~ spl3_57), inference(avatar_component_clause, [], [f529])).
fof(f220, plain, ~ (op(e0, e0) = op(e0, e1)), inference(cnf_transformation, [], [f9])).
fof(f2956, plain, (~ spl3_4 | ~ spl3_36), inference(avatar_split_clause, [], [f2955, f439, f303])).
fof(f303, plain, (spl3_4 <=> (e3 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl3_4])])).
fof(f439, plain, (spl3_36 <=> (e3 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl3_36])])).
fof(f2955, plain, (~ (e3 = op(e3, e3)) | ~ spl3_36), inference(forward_demodulation, [], [f218, f441])).
fof(f441, plain, ((e3 = op(e1, e3)) | ~ spl3_36), inference(avatar_component_clause, [], [f439])).
fof(f218, plain, ~ (op(e1, e3) = op(e3, e3)), inference(cnf_transformation, [], [f9])).
fof(f2787, plain, (~ spl3_53 | ~ spl3_57), inference(avatar_split_clause, [], [f2766, f529, f512])).
fof(f512, plain, (spl3_53 <=> (e0 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl3_53])])).
fof(f2766, plain, (~ (e0 = op(e0, e2)) | ~ spl3_57), inference(backward_demodulation, [], [f223, f531])).
fof(f223, plain, ~ (op(e0, e1) = op(e0, e2)), inference(cnf_transformation, [], [f9])).
fof(f2752, plain, (spl3_19 | ~ spl3_68), inference(avatar_split_clause, [], [f2728, f575, f367])).
fof(f575, plain, (spl3_68 <=> (e3 = unit)), introduced(avatar_definition, [new_symbols(naming, [spl3_68])])).
fof(f2728, plain, ((e2 = op(e2, e3)) | ~ spl3_68), inference(backward_demodulation, [], [f153, f577])).
fof(f577, plain, ((e3 = unit) | ~ spl3_68), inference(avatar_component_clause, [], [f575])).
fof(f2751, plain, (spl3_7 | ~ spl3_68), inference(avatar_contradiction_clause, [], [f2750])).
fof(f2750, plain, ($false | (spl3_7 | ~ spl3_68)), inference(subsumption_resolution, [], [f2727, f317])).
fof(f317, plain, (~ (e2 = op(e3, e2)) | spl3_7), inference(avatar_component_clause, [], [f316])).
fof(f316, plain, (spl3_7 <=> (e2 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl3_7])])).
fof(f2727, plain, ((e2 = op(e3, e2)) | ~ spl3_68), inference(backward_demodulation, [], [f152, f577])).
fof(f152, plain, (e2 = op(unit, e2)), inference(cnf_transformation, [], [f3])).
fof(f2747, plain, (spl3_10 | ~ spl3_68), inference(avatar_split_clause, [], [f2725, f575, f329])).
fof(f329, plain, (spl3_10 <=> (e1 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl3_10])])).
fof(f2725, plain, ((e1 = op(e3, e1)) | ~ spl3_68), inference(backward_demodulation, [], [f150, f577])).
fof(f150, plain, (e1 = op(unit, e1)), inference(cnf_transformation, [], [f3])).
fof(f2744, plain, (spl3_13 | ~ spl3_68), inference(avatar_split_clause, [], [f2723, f575, f342])).
fof(f2723, plain, ((e0 = op(e3, e0)) | ~ spl3_68), inference(backward_demodulation, [], [f148, f577])).
fof(f2634, plain, (~ spl3_10 | ~ spl3_42), inference(avatar_split_clause, [], [f2633, f465, f329])).
fof(f465, plain, (spl3_42 <=> (e1 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl3_42])])).
fof(f2633, plain, (~ (e1 = op(e3, e1)) | ~ spl3_42), inference(forward_demodulation, [], [f206, f467])).
fof(f467, plain, ((e1 = op(e1, e1)) | ~ spl3_42), inference(avatar_component_clause, [], [f465])).
fof(f206, plain, ~ (op(e1, e1) = op(e3, e1)), inference(cnf_transformation, [], [f9])).
fof(f2598, plain, (~ spl3_46 | ~ spl3_62), inference(avatar_split_clause, [], [f2597, f550, f482])).
fof(f482, plain, (spl3_46 <=> (e1 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl3_46])])).
fof(f550, plain, (spl3_62 <=> (op(e0, e0) = e1)), introduced(avatar_definition, [new_symbols(naming, [spl3_62])])).
fof(f2597, plain, (~ (e1 = op(e1, e0)) | ~ spl3_62), inference(forward_demodulation, [], [f196, f552])).
fof(f552, plain, ((op(e0, e0) = e1) | ~ spl3_62), inference(avatar_component_clause, [], [f550])).
fof(f196, plain, ~ (op(e0, e0) = op(e1, e0)), inference(cnf_transformation, [], [f9])).
fof(f2568, plain, (~ spl3_30 | ~ spl3_62), inference(avatar_split_clause, [], [f2567, f550, f414])).
fof(f414, plain, (spl3_30 <=> (e1 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl3_30])])).
fof(f2567, plain, (~ (e1 = op(e2, e0)) | ~ spl3_62), inference(forward_demodulation, [], [f197, f552])).
fof(f197, plain, ~ (op(e0, e0) = op(e2, e0)), inference(cnf_transformation, [], [f9])).
fof(f2556, plain, (~ spl3_26 | ~ spl3_42), inference(avatar_split_clause, [], [f2555, f465, f397])).
fof(f397, plain, (spl3_26 <=> (e1 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl3_26])])).
fof(f2555, plain, (~ (e1 = op(e2, e1)) | ~ spl3_42), inference(forward_demodulation, [], [f205, f467])).
fof(f205, plain, ~ (op(e1, e1) = op(e2, e1)), inference(cnf_transformation, [], [f9])).
fof(f2291, plain, (~ spl3_61 | ~ spl3_53), inference(avatar_split_clause, [], [f2290, f512, f546])).
fof(f2290, plain, (~ (e0 = op(e0, e0)) | ~ spl3_53), inference(forward_demodulation, [], [f221, f514])).
fof(f514, plain, ((e0 = op(e0, e2)) | ~ spl3_53), inference(avatar_component_clause, [], [f512])).
fof(f221, plain, ~ (op(e0, e0) = op(e0, e2)), inference(cnf_transformation, [], [f9])).
fof(f2260, plain, (~ spl3_52 | ~ spl3_20), inference(avatar_split_clause, [], [f2259, f371, f507])).
fof(f507, plain, (spl3_52 <=> (e3 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl3_52])])).
fof(f371, plain, (spl3_20 <=> (e3 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl3_20])])).
fof(f2259, plain, (~ (e3 = op(e0, e3)) | ~ spl3_20), inference(forward_demodulation, [], [f215, f373])).
fof(f373, plain, ((e3 = op(e2, e3)) | ~ spl3_20), inference(avatar_component_clause, [], [f371])).
fof(f215, plain, ~ (op(e0, e3) = op(e2, e3)), inference(cnf_transformation, [], [f9])).
fof(f2221, plain, (~ spl3_10 | ~ spl3_26), inference(avatar_split_clause, [], [f2208, f397, f329])).
fof(f2208, plain, (~ (e1 = op(e3, e1)) | ~ spl3_26), inference(backward_demodulation, [], [f207, f399])).
fof(f399, plain, ((e1 = op(e2, e1)) | ~ spl3_26), inference(avatar_component_clause, [], [f397])).
fof(f207, plain, ~ (op(e2, e1) = op(e3, e1)), inference(cnf_transformation, [], [f9])).
fof(f2147, plain, (~ spl3_59 | ~ spl3_63), inference(avatar_contradiction_clause, [], [f2146])).
fof(f2146, plain, ($false | (~ spl3_59 | ~ spl3_63)), inference(subsumption_resolution, [], [f2133, f2054])).
fof(f2054, plain, (~ (op(e0, e2) = op(e2, e1)) | ~ spl3_63), inference(backward_demodulation, [], [f232, f2048])).
fof(f2048, plain, ((op(e0, e2) = op(e2, e0)) | ~ spl3_63), inference(forward_demodulation, [], [f84, f556])).
fof(f556, plain, ((op(e0, e0) = e2) | ~ spl3_63), inference(avatar_component_clause, [], [f554])).
fof(f554, plain, (spl3_63 <=> (op(e0, e0) = e2)), introduced(avatar_definition, [new_symbols(naming, [spl3_63])])).
fof(f84, plain, (op(op(e0, e0), e0) = op(e0, op(e0, e0))), inference(cnf_transformation, [], [f2])).
fof(f2, plain, ((op(op(e3, e3), e3) = op(e3, op(e3, e3))) & (op(op(e3, e3), e2) = op(e3, op(e3, e2))) & (op(op(e3, e3), e1) = op(e3, op(e3, e1))) & (op(op(e3, e3), e0) = op(e3, op(e3, e0))) & (op(op(e3, e2), e3) = op(e3, op(e2, e3))) & (op(op(e3, e2), e2) = op(e3, op(e2, e2))) & (op(op(e3, e2), e1) = op(e3, op(e2, e1))) & (op(op(e3, e2), e0) = op(e3, op(e2, e0))) & (op(op(e3, e1), e3) = op(e3, op(e1, e3))) & (op(op(e3, e1), e2) = op(e3, op(e1, e2))) & (op(op(e3, e1), e1) = op(e3, op(e1, e1))) & (op(op(e3, e1), e0) = op(e3, op(e1, e0))) & (op(op(e3, e0), e3) = op(e3, op(e0, e3))) & (op(op(e3, e0), e2) = op(e3, op(e0, e2))) & (op(op(e3, e0), e1) = op(e3, op(e0, e1))) & (op(op(e3, e0), e0) = op(e3, op(e0, e0))) & (op(op(e2, e3), e3) = op(e2, op(e3, e3))) & (op(op(e2, e3), e2) = op(e2, op(e3, e2))) & (op(op(e2, e3), e1) = op(e2, op(e3, e1))) & (op(op(e2, e3), e0) = op(e2, op(e3, e0))) & (op(op(e2, e2), e3) = op(e2, op(e2, e3))) & (op(op(e2, e2), e2) = op(e2, op(e2, e2))) & (op(op(e2, e2), e1) = op(e2, op(e2, e1))) & (op(op(e2, e2), e0) = op(e2, op(e2, e0))) & (op(op(e2, e1), e3) = op(e2, op(e1, e3))) & (op(op(e2, e1), e2) = op(e2, op(e1, e2))) & (op(op(e2, e1), e1) = op(e2, op(e1, e1))) & (op(op(e2, e1), e0) = op(e2, op(e1, e0))) & (op(op(e2, e0), e3) = op(e2, op(e0, e3))) & (op(op(e2, e0), e2) = op(e2, op(e0, e2))) & (op(op(e2, e0), e1) = op(e2, op(e0, e1))) & (op(op(e2, e0), e0) = op(e2, op(e0, e0))) & (op(op(e1, e3), e3) = op(e1, op(e3, e3))) & (op(op(e1, e3), e2) = op(e1, op(e3, e2))) & (op(op(e1, e3), e1) = op(e1, op(e3, e1))) & (op(op(e1, e3), e0) = op(e1, op(e3, e0))) & (op(op(e1, e2), e3) = op(e1, op(e2, e3))) & (op(op(e1, e2), e2) = op(e1, op(e2, e2))) & (op(op(e1, e2), e1) = op(e1, op(e2, e1))) & (op(op(e1, e2), e0) = op(e1, op(e2, e0))) & (op(op(e1, e1), e3) = op(e1, op(e1, e3))) & (op(op(e1, e1), e2) = op(e1, op(e1, e2))) & (op(op(e1, e1), e1) = op(e1, op(e1, e1))) & (op(op(e1, e1), e0) = op(e1, op(e1, e0))) & (op(op(e1, e0), e3) = op(e1, op(e0, e3))) & (op(op(e1, e0), e2) = op(e1, op(e0, e2))) & (op(op(e1, e0), e1) = op(e1, op(e0, e1))) & (op(op(e1, e0), e0) = op(e1, op(e0, e0))) & (op(op(e0, e3), e3) = op(e0, op(e3, e3))) & (op(op(e0, e3), e2) = op(e0, op(e3, e2))) & (op(op(e0, e3), e1) = op(e0, op(e3, e1))) & (op(op(e0, e3), e0) = op(e0, op(e3, e0))) & (op(op(e0, e2), e3) = op(e0, op(e2, e3))) & (op(op(e0, e2), e2) = op(e0, op(e2, e2))) & (op(op(e0, e2), e1) = op(e0, op(e2, e1))) & (op(op(e0, e2), e0) = op(e0, op(e2, e0))) & (op(op(e0, e1), e3) = op(e0, op(e1, e3))) & (op(op(e0, e1), e2) = op(e0, op(e1, e2))) & (op(op(e0, e1), e1) = op(e0, op(e1, e1))) & (op(op(e0, e1), e0) = op(e0, op(e1, e0))) & (op(op(e0, e0), e3) = op(e0, op(e0, e3))) & (op(op(e0, e0), e2) = op(e0, op(e0, e2))) & (op(op(e0, e0), e1) = op(e0, op(e0, e1))) & (op(op(e0, e0), e0) = op(e0, op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG023+1.p', ax2)).
fof(f232, plain, ~ (op(e2, e0) = op(e2, e1)), inference(cnf_transformation, [], [f9])).
fof(f2133, plain, ((op(e0, e2) = op(e2, e1)) | (~ spl3_59 | ~ spl3_63)), inference(backward_demodulation, [], [f2063, f539])).
fof(f539, plain, ((e2 = op(e0, e1)) | ~ spl3_59), inference(avatar_component_clause, [], [f537])).
fof(f537, plain, (spl3_59 <=> (e2 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl3_59])])).
fof(f2063, plain, ((op(e2, e1) = op(e0, op(e0, e1))) | ~ spl3_63), inference(forward_demodulation, [], [f85, f556])).
fof(f85, plain, (op(op(e0, e0), e1) = op(e0, op(e0, e1))), inference(cnf_transformation, [], [f2])).
fof(f2076, plain, (~ spl3_31 | ~ spl3_63), inference(avatar_split_clause, [], [f2075, f554, f418])).
fof(f418, plain, (spl3_31 <=> (e2 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl3_31])])).
fof(f2075, plain, (~ (e2 = op(e2, e0)) | ~ spl3_63), inference(forward_demodulation, [], [f197, f556])).
fof(f2027, plain, (spl3_53 | ~ spl3_67), inference(avatar_split_clause, [], [f1765, f571, f512])).
fof(f571, plain, (spl3_67 <=> (e2 = unit)), introduced(avatar_definition, [new_symbols(naming, [spl3_67])])).
fof(f1765, plain, ((e0 = op(e0, e2)) | ~ spl3_67), inference(backward_demodulation, [], [f149, f573])).
fof(f573, plain, ((e2 = unit) | ~ spl3_67), inference(avatar_component_clause, [], [f571])).
fof(f149, plain, (e0 = op(e0, unit)), inference(cnf_transformation, [], [f3])).
fof(f2013, plain, (spl3_26 | ~ spl3_67), inference(avatar_split_clause, [], [f1766, f571, f397])).
fof(f1766, plain, ((e1 = op(e2, e1)) | ~ spl3_67), inference(backward_demodulation, [], [f150, f573])).
fof(f1972, plain, (~ spl3_7 | ~ spl3_3), inference(avatar_split_clause, [], [f1971, f299, f316])).
fof(f299, plain, (spl3_3 <=> (e2 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl3_3])])).
fof(f1971, plain, (~ (e2 = op(e3, e2)) | ~ spl3_3), inference(forward_demodulation, [], [f243, f301])).
fof(f301, plain, ((e2 = op(e3, e3)) | ~ spl3_3), inference(avatar_component_clause, [], [f299])).
fof(f243, plain, ~ (op(e3, e2) = op(e3, e3)), inference(cnf_transformation, [], [f9])).
fof(f1938, plain, (~ spl3_23 | ~ spl3_24), inference(avatar_contradiction_clause, [], [f1937])).
fof(f1937, plain, ($false | (~ spl3_23 | ~ spl3_24)), inference(subsumption_resolution, [], [f1936, f249])).
fof(f249, plain, ~ (e2 = e3), inference(cnf_transformation, [], [f10])).
fof(f10, plain, (~ (e2 = e3) & ~ (e1 = e3) & ~ (e1 = e2) & ~ (e0 = e3) & ~ (e0 = e2) & ~ (e0 = e1)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG023+1.p', ax10)).
fof(f1936, plain, ((e2 = e3) | (~ spl3_23 | ~ spl3_24)), inference(backward_demodulation, [], [f390, f386])).
fof(f386, plain, ((e2 = op(e2, e2)) | ~ spl3_23), inference(avatar_component_clause, [], [f384])).
fof(f390, plain, ((e3 = op(e2, e2)) | ~ spl3_24), inference(avatar_component_clause, [], [f388])).
fof(f388, plain, (spl3_24 <=> (e3 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl3_24])])).
fof(f1795, plain, (spl3_23 | ~ spl3_67), inference(avatar_split_clause, [], [f1769, f571, f384])).
fof(f1769, plain, ((e2 = op(e2, e2)) | ~ spl3_67), inference(backward_demodulation, [], [f153, f573])).
fof(f1663, plain, (~ spl3_16 | ~ spl3_64), inference(avatar_split_clause, [], [f1662, f558, f354])).
fof(f354, plain, (spl3_16 <=> (e3 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl3_16])])).
fof(f558, plain, (spl3_64 <=> (op(e0, e0) = e3)), introduced(avatar_definition, [new_symbols(naming, [spl3_64])])).
fof(f1662, plain, (~ (e3 = op(e3, e0)) | ~ spl3_64), inference(forward_demodulation, [], [f198, f560])).
fof(f560, plain, ((op(e0, e0) = e3) | ~ spl3_64), inference(avatar_component_clause, [], [f558])).
fof(f1614, plain, (~ spl3_24 | ~ spl3_28), inference(avatar_contradiction_clause, [], [f1613])).
fof(f1613, plain, ($false | (~ spl3_24 | ~ spl3_28)), inference(subsumption_resolution, [], [f1612, f407])).
fof(f407, plain, ((e3 = op(e2, e1)) | ~ spl3_28), inference(avatar_component_clause, [], [f405])).
fof(f405, plain, (spl3_28 <=> (e3 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl3_28])])).
fof(f1612, plain, (~ (e3 = op(e2, e1)) | ~ spl3_24), inference(forward_demodulation, [], [f235, f390])).
fof(f1254, plain, (~ spl3_19 | ~ spl3_31), inference(avatar_split_clause, [], [f1243, f418, f367])).
fof(f1243, plain, (~ (e2 = op(e2, e3)) | ~ spl3_31), inference(backward_demodulation, [], [f234, f420])).
fof(f420, plain, ((e2 = op(e2, e0)) | ~ spl3_31), inference(avatar_component_clause, [], [f418])).
fof(f234, plain, ~ (op(e2, e0) = op(e2, e3)), inference(cnf_transformation, [], [f9])).
fof(f1253, plain, (~ spl3_27 | ~ spl3_31), inference(avatar_split_clause, [], [f1242, f418, f401])).
fof(f1242, plain, (~ (e2 = op(e2, e1)) | ~ spl3_31), inference(backward_demodulation, [], [f232, f420])).
fof(f1176, plain, (~ spl3_26 | ~ spl3_58), inference(avatar_split_clause, [], [f1167, f533, f397])).
fof(f533, plain, (spl3_58 <=> (e1 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl3_58])])).
fof(f1167, plain, (~ (e1 = op(e2, e1)) | ~ spl3_58), inference(backward_demodulation, [], [f203, f535])).
fof(f535, plain, ((e1 = op(e0, e1)) | ~ spl3_58), inference(avatar_component_clause, [], [f533])).
fof(f203, plain, ~ (op(e0, e1) = op(e2, e1)), inference(cnf_transformation, [], [f9])).
fof(f1151, plain, (spl3_36 | ~ spl3_66), inference(avatar_split_clause, [], [f1132, f567, f439])).
fof(f1132, plain, ((e3 = op(e1, e3)) | ~ spl3_66), inference(backward_demodulation, [], [f154, f569])).
fof(f154, plain, (e3 = op(unit, e3)), inference(cnf_transformation, [], [f3])).
fof(f1148, plain, (spl3_42 | ~ spl3_66), inference(avatar_split_clause, [], [f1129, f567, f465])).
fof(f1129, plain, ((e1 = op(e1, e1)) | ~ spl3_66), inference(backward_demodulation, [], [f151, f569])).
fof(f151, plain, (e1 = op(e1, unit)), inference(cnf_transformation, [], [f3])).
fof(f1146, plain, (spl3_57 | ~ spl3_66), inference(avatar_split_clause, [], [f1127, f567, f529])).
fof(f1127, plain, ((e0 = op(e0, e1)) | ~ spl3_66), inference(backward_demodulation, [], [f149, f569])).
fof(f940, plain, (~ spl3_29 | ~ spl3_45), inference(avatar_split_clause, [], [f929, f478, f410])).
fof(f410, plain, (spl3_29 <=> (e0 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl3_29])])).
fof(f929, plain, (~ (e0 = op(e2, e0)) | ~ spl3_45), inference(backward_demodulation, [], [f199, f480])).
fof(f480, plain, ((e0 = op(e1, e0)) | ~ spl3_45), inference(avatar_component_clause, [], [f478])).
fof(f199, plain, ~ (op(e1, e0) = op(e2, e0)), inference(cnf_transformation, [], [f9])).
fof(f853, plain, (spl3_16 | ~ spl3_65), inference(avatar_split_clause, [], [f828, f563, f354])).
fof(f563, plain, (spl3_65 <=> (e0 = unit)), introduced(avatar_definition, [new_symbols(naming, [spl3_65])])).
fof(f828, plain, ((e3 = op(e3, e0)) | ~ spl3_65), inference(backward_demodulation, [], [f155, f565])).
fof(f565, plain, ((e0 = unit) | ~ spl3_65), inference(avatar_component_clause, [], [f563])).
fof(f155, plain, (e3 = op(e3, unit)), inference(cnf_transformation, [], [f3])).
fof(f852, plain, (spl3_52 | ~ spl3_65), inference(avatar_split_clause, [], [f827, f563, f507])).
fof(f827, plain, ((e3 = op(e0, e3)) | ~ spl3_65), inference(backward_demodulation, [], [f154, f565])).
fof(f851, plain, (spl3_31 | ~ spl3_65), inference(avatar_split_clause, [], [f826, f563, f418])).
fof(f826, plain, ((e2 = op(e2, e0)) | ~ spl3_65), inference(backward_demodulation, [], [f153, f565])).
fof(f849, plain, (spl3_46 | ~ spl3_65), inference(avatar_split_clause, [], [f824, f563, f482])).
fof(f824, plain, ((e1 = op(e1, e0)) | ~ spl3_65), inference(backward_demodulation, [], [f151, f565])).
fof(f848, plain, (spl3_58 | ~ spl3_65), inference(avatar_split_clause, [], [f823, f563, f533])).
fof(f823, plain, ((e1 = op(e0, e1)) | ~ spl3_65), inference(backward_demodulation, [], [f150, f565])).
fof(f706, plain, (~ spl3_44 | ~ spl3_25), inference(avatar_split_clause, [], [f285, f393, f473])).
fof(f473, plain, (spl3_44 <=> (e3 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl3_44])])).
fof(f393, plain, (spl3_25 <=> (e0 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl3_25])])).
fof(f285, plain, (~ (e0 = op(e2, e1)) | ~ (e3 = op(e1, e1))), inference(cnf_transformation, [], [f56])).
fof(f56, plain, (~ (e0 = op(e2, e1)) | ~ (e3 = op(e1, e1))), inference(ennf_transformation, [], [f31])).
fof(f31, plain, ~ ((e0 = op(e2, e1)) & (e3 = op(e1, e1))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG023+1.p', ax31)).
fof(f703, plain, (~ spl3_43 | ~ spl3_60), inference(avatar_split_clause, [], [f282, f541, f469])).
fof(f469, plain, (spl3_43 <=> (e2 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl3_43])])).
fof(f541, plain, (spl3_60 <=> (e3 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl3_60])])).
fof(f282, plain, (~ (e3 = op(e0, e1)) | ~ (e2 = op(e1, e1))), inference(cnf_transformation, [], [f53])).
fof(f53, plain, (~ (e3 = op(e0, e1)) | ~ (e2 = op(e1, e1))), inference(ennf_transformation, [], [f28])).
fof(f28, plain, ~ ((e3 = op(e0, e1)) & (e2 = op(e1, e1))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG023+1.p', ax28)).
fof(f698, plain, (~ spl3_62 | ~ spl3_32), inference(avatar_split_clause, [], [f277, f422, f550])).
fof(f422, plain, (spl3_32 <=> (e3 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl3_32])])).
fof(f277, plain, (~ (e3 = op(e2, e0)) | ~ (op(e0, e0) = e1)), inference(cnf_transformation, [], [f48])).
fof(f48, plain, (~ (e3 = op(e2, e0)) | ~ (op(e0, e0) = e1)), inference(ennf_transformation, [], [f23])).
fof(f23, plain, ~ ((e3 = op(e2, e0)) & (op(e0, e0) = e1)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG023+1.p', ax23)).
fof(f687, plain, (~ spl3_1 | ~ spl3_18), inference(avatar_split_clause, [], [f266, f363, f291])).
fof(f291, plain, (spl3_1 <=> (e0 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl3_1])])).
fof(f363, plain, (spl3_18 <=> (e1 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl3_18])])).
fof(f266, plain, (~ (e1 = op(e2, e3)) | ~ (e0 = op(e3, e3))), inference(cnf_transformation, [], [f37])).
fof(f37, plain, (~ (e1 = op(e2, e3)) | ~ (e0 = op(e3, e3))), inference(ennf_transformation, [], [f12])).
fof(f12, plain, ~ ((e1 = op(e2, e3)) & (e0 = op(e3, e3))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG023+1.p', ax12)).
fof(f686, plain, (spl3_87 | spl3_86 | spl3_85 | spl3_64), inference(avatar_split_clause, [], [f262, f558, f660, f668, f676])).
fof(f676, plain, (spl3_87 <=> sP0), introduced(avatar_definition, [new_symbols(naming, [spl3_87])])).
fof(f668, plain, (spl3_86 <=> sP1), introduced(avatar_definition, [new_symbols(naming, [spl3_86])])).
fof(f660, plain, (spl3_85 <=> sP2), introduced(avatar_definition, [new_symbols(naming, [spl3_85])])).
fof(f262, plain, ((op(e0, e0) = e3) | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f64])).
fof(f64, plain, (((e3 = op(e3, e3)) & (e3 = op(e2, e2)) & (e3 = op(e1, e1)) & (op(e0, e0) = e3)) | sP2 | sP1 | sP0), inference(definition_folding, [], [f11, e63, e62, e61])).
fof(f61, plain, (((e0 = op(e3, e3)) & (e0 = op(e2, e2)) & (e0 = op(e1, e1)) & (e0 = op(e0, e0))) | ~ sP0), inference(usedef, [], [e61])).
fof(e61, plain, (sP0 <=> ((e0 = op(e3, e3)) & (e0 = op(e2, e2)) & (e0 = op(e1, e1)) & (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f62, plain, (((e1 = op(e3, e3)) & (e1 = op(e2, e2)) & (e1 = op(e1, e1)) & (op(e0, e0) = e1)) | ~ sP1), inference(usedef, [], [e62])).
fof(e62, plain, (sP1 <=> ((e1 = op(e3, e3)) & (e1 = op(e2, e2)) & (e1 = op(e1, e1)) & (op(e0, e0) = e1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f63, plain, (((e2 = op(e3, e3)) & (e2 = op(e2, e2)) & (e2 = op(e1, e1)) & (op(e0, e0) = e2)) | ~ sP2), inference(usedef, [], [e63])).
fof(e63, plain, (sP2 <=> ((e2 = op(e3, e3)) & (e2 = op(e2, e2)) & (e2 = op(e1, e1)) & (op(e0, e0) = e2))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f11, plain, (((e3 = op(e3, e3)) & (e3 = op(e2, e2)) & (e3 = op(e1, e1)) & (op(e0, e0) = e3)) | ((e2 = op(e3, e3)) & (e2 = op(e2, e2)) & (e2 = op(e1, e1)) & (op(e0, e0) = e2)) | ((e1 = op(e3, e3)) & (e1 = op(e2, e2)) & (e1 = op(e1, e1)) & (op(e0, e0) = e1)) | ((e0 = op(e3, e3)) & (e0 = op(e2, e2)) & (e0 = op(e1, e1)) & (e0 = op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG023+1.p', ax11)).
fof(f685, plain, (spl3_87 | spl3_86 | spl3_85 | spl3_44), inference(avatar_split_clause, [], [f263, f473, f660, f668, f676])).
fof(f263, plain, ((e3 = op(e1, e1)) | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f64])).
fof(f684, plain, (spl3_87 | spl3_86 | spl3_85 | spl3_24), inference(avatar_split_clause, [], [f264, f388, f660, f668, f676])).
fof(f264, plain, ((e3 = op(e2, e2)) | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f64])).
fof(f683, plain, (spl3_87 | spl3_86 | spl3_85 | spl3_4), inference(avatar_split_clause, [], [f265, f303, f660, f668, f676])).
fof(f265, plain, ((e3 = op(e3, e3)) | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f64])).
fof(f682, plain, (~ spl3_87 | spl3_61), inference(avatar_split_clause, [], [f258, f546, f676])).
fof(f258, plain, ((e0 = op(e0, e0)) | ~ sP0), inference(cnf_transformation, [], [f67])).
fof(f67, plain, (((e0 = op(e3, e3)) & (e0 = op(e2, e2)) & (e0 = op(e1, e1)) & (e0 = op(e0, e0))) | ~ sP0), inference(nnf_transformation, [], [f61])).
fof(f680, plain, (~ spl3_87 | spl3_21), inference(avatar_split_clause, [], [f260, f376, f676])).
fof(f260, plain, ((e0 = op(e2, e2)) | ~ sP0), inference(cnf_transformation, [], [f67])).
fof(f679, plain, (~ spl3_87 | spl3_1), inference(avatar_split_clause, [], [f261, f291, f676])).
fof(f261, plain, ((e0 = op(e3, e3)) | ~ sP0), inference(cnf_transformation, [], [f67])).
fof(f674, plain, (~ spl3_86 | spl3_62), inference(avatar_split_clause, [], [f254, f550, f668])).
fof(f254, plain, ((op(e0, e0) = e1) | ~ sP1), inference(cnf_transformation, [], [f66])).
fof(f66, plain, (((e1 = op(e3, e3)) & (e1 = op(e2, e2)) & (e1 = op(e1, e1)) & (op(e0, e0) = e1)) | ~ sP1), inference(nnf_transformation, [], [f62])).
fof(f673, plain, (~ spl3_86 | spl3_42), inference(avatar_split_clause, [], [f255, f465, f668])).
fof(f255, plain, ((e1 = op(e1, e1)) | ~ sP1), inference(cnf_transformation, [], [f66])).
fof(f666, plain, (~ spl3_85 | spl3_63), inference(avatar_split_clause, [], [f250, f554, f660])).
fof(f250, plain, ((op(e0, e0) = e2) | ~ sP2), inference(cnf_transformation, [], [f65])).
fof(f65, plain, (((e2 = op(e3, e3)) & (e2 = op(e2, e2)) & (e2 = op(e1, e1)) & (op(e0, e0) = e2)) | ~ sP2), inference(nnf_transformation, [], [f63])).
fof(f665, plain, (~ spl3_85 | spl3_43), inference(avatar_split_clause, [], [f251, f469, f660])).
fof(f251, plain, ((e2 = op(e1, e1)) | ~ sP2), inference(cnf_transformation, [], [f65])).
fof(f664, plain, (~ spl3_85 | spl3_23), inference(avatar_split_clause, [], [f252, f384, f660])).
fof(f252, plain, ((e2 = op(e2, e2)) | ~ sP2), inference(cnf_transformation, [], [f65])).
fof(f663, plain, (~ spl3_85 | spl3_3), inference(avatar_split_clause, [], [f253, f299, f660])).
fof(f253, plain, ((e2 = op(e3, e3)) | ~ sP2), inference(cnf_transformation, [], [f65])).
fof(f578, plain, (spl3_65 | spl3_66 | spl3_67 | spl3_68), inference(avatar_split_clause, [], [f156, f575, f571, f567, f563])).
fof(f156, plain, ((e3 = unit) | (e2 = unit) | (e1 = unit) | (e0 = unit)), inference(cnf_transformation, [], [f3])).
fof(f544, plain, (spl3_57 | spl3_58 | spl3_59 | spl3_60), inference(avatar_split_clause, [], [f69, f541, f537, f533, f529])).
fof(f69, plain, ((e3 = op(e0, e1)) | (e2 = op(e0, e1)) | (e1 = op(e0, e1)) | (e0 = op(e0, e1))), inference(cnf_transformation, [], [f1])).
fof(f1, plain, (((e3 = op(e3, e3)) | (e2 = op(e3, e3)) | (e1 = op(e3, e3)) | (e0 = op(e3, e3))) & ((e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (e0 = op(e3, e2))) & ((e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (e0 = op(e3, e1))) & ((e3 = op(e3, e0)) | (e2 = op(e3, e0)) | (e1 = op(e3, e0)) | (e0 = op(e3, e0))) & ((e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))) & ((e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (e0 = op(e2, e2))) & ((e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))) & ((e3 = op(e2, e0)) | (e2 = op(e2, e0)) | (e1 = op(e2, e0)) | (e0 = op(e2, e0))) & ((e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (e0 = op(e1, e3))) & ((e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (e0 = op(e1, e2))) & ((e3 = op(e1, e1)) | (e2 = op(e1, e1)) | (e1 = op(e1, e1)) | (e0 = op(e1, e1))) & ((e3 = op(e1, e0)) | (e2 = op(e1, e0)) | (e1 = op(e1, e0)) | (e0 = op(e1, e0))) & ((e3 = op(e0, e3)) | (e2 = op(e0, e3)) | (e1 = op(e0, e3)) | (e0 = op(e0, e3))) & ((e3 = op(e0, e2)) | (e2 = op(e0, e2)) | (e1 = op(e0, e2)) | (e0 = op(e0, e2))) & ((e3 = op(e0, e1)) | (e2 = op(e0, e1)) | (e1 = op(e0, e1)) | (e0 = op(e0, e1))) & ((op(e0, e0) = e3) | (op(e0, e0) = e2) | (op(e0, e0) = e1) | (e0 = op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG023+1.p', ax1)).
fof(f425, plain, (spl3_29 | spl3_30 | spl3_31 | spl3_32), inference(avatar_split_clause, [], [f76, f422, f418, f414, f410])).
fof(f76, plain, ((e3 = op(e2, e0)) | (e2 = op(e2, e0)) | (e1 = op(e2, e0)) | (e0 = op(e2, e0))), inference(cnf_transformation, [], [f1])).
fof(f408, plain, (spl3_25 | spl3_26 | spl3_27 | spl3_28), inference(avatar_split_clause, [], [f77, f405, f401, f397, f393])).
fof(f77, plain, ((e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))), inference(cnf_transformation, [], [f1])).
fof(f374, plain, (spl3_17 | spl3_18 | spl3_19 | spl3_20), inference(avatar_split_clause, [], [f79, f371, f367, f363, f359])).
fof(f79, plain, ((e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))), inference(cnf_transformation, [], [f1])).