fof(f5834, plain, $false, inference(avatar_sat_refutation, [], [f801, f920, f923, f925, f926, f937, f1241, f1244, f1246, f1247, f1258, f1259, f1278, f1288, f1294, f1295, f1297, f1300, f1308, f1316, f1322, f1323, f1325, f1328, f1330, f1331, f1332, f1333, f1438, f1459, f2084, f3432, f3433, f3477, f3492, f3549, f3554, f3591, f3593, f3665, f3727, f3730, f3780, f3784, f3790, f3814, f3846, f3873, f3881, f3899, f3922, f3926, f3953, f3999, f4071, f4100, f4137, f4143, f4183, f4209, f4249, f4260, f4280, f4438, f4453, f4494, f4507, f4512, f4525, f4547, f4606, f4664, f4680, f4706, f4766, f4837, f4873, f4874, f4875, f4876, f4910, f4911, f4912, f4915, f4916, f4934, f4947, f5039, f5042, f5046, f5102, f5175, f5226, f5288, f5346, f5402, f5460, f5527, f5588, f5650, f5708, f5768, f5822])).
fof(f5822, plain, (~ spl42_61 | ~ spl42_129 | ~ spl42_148 | spl42_243), inference(avatar_contradiction_clause, [], [f5821])).
fof(f5821, plain, ($false | (~ spl42_61 | ~ spl42_129 | ~ spl42_148 | spl42_243)), inference(subsumption_resolution, [], [f5820, f1211])).
fof(f1211, plain, ((e20 = op2(e20, e20)) | ~ spl42_129), inference(avatar_component_clause, [], [f1209])).
fof(f1209, plain, (spl42_129 <=> (e20 = op2(e20, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_129])])).
fof(f5820, plain, (~ (e20 = op2(e20, e20)) | (~ spl42_61 | ~ spl42_148 | spl42_243)), inference(forward_demodulation, [], [f5819, f4429])).
fof(f4429, plain, ((e20 = h11(e10)) | ~ spl42_148), inference(backward_demodulation, [], [f1392, f1426])).
fof(f1426, plain, ((e20 = h10(e10)) | ~ spl42_148), inference(avatar_component_clause, [], [f1425])).
fof(f1425, plain, (spl42_148 <=> (e20 = h10(e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_148])])).
fof(f1392, plain, (h10(e10) = h11(e10)), inference(forward_demodulation, [], [f434, f430])).
fof(f430, plain, (op2(e23, e23) = h10(e10)), inference(cnf_transformation, [], [f25])).
fof(f25, plain, ((op2(e20, e23) = h10(e12)) & (op2(e23, e23) = h10(e10)) & (e20 = h10(e13)) & (e23 = h10(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG037+1.p', ax25)).
fof(f434, plain, (op2(e23, e23) = h11(e10)), inference(cnf_transformation, [], [f26])).
fof(f26, plain, ((op2(e21, e23) = h11(e12)) & (op2(e23, e23) = h11(e10)) & (e21 = h11(e13)) & (e23 = h11(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG037+1.p', ax26)).
fof(f5819, plain, (~ (op2(e20, e20) = h11(e10)) | (~ spl42_61 | ~ spl42_148 | spl42_243)), inference(forward_demodulation, [], [f5818, f1426])).
fof(f5818, plain, (~ (h11(e10) = op2(h10(e10), h10(e10))) | (~ spl42_61 | spl42_243)), inference(forward_demodulation, [], [f2055, f890])).
fof(f890, plain, ((e10 = op1(e10, e10)) | ~ spl42_61), inference(avatar_component_clause, [], [f888])).
fof(f888, plain, (spl42_61 <=> (e10 = op1(e10, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_61])])).
fof(f2055, plain, (~ (op2(h10(e10), h10(e10)) = h11(op1(e10, e10))) | spl42_243), inference(avatar_component_clause, [], [f2053])).
fof(f2053, plain, (spl42_243 <=> (op2(h10(e10), h10(e10)) = h11(op1(e10, e10)))), introduced(avatar_definition, [new_symbols(naming, [spl42_243])])).
fof(f5768, plain, (~ spl42_28 | ~ spl42_86 | ~ spl42_150 | spl42_234), inference(avatar_contradiction_clause, [], [f5767])).
fof(f5767, plain, ($false | (~ spl42_28 | ~ spl42_86 | ~ spl42_150 | spl42_234)), inference(subsumption_resolution, [], [f5766, f1028])).
fof(f1028, plain, ((e21 = op2(e22, e23)) | ~ spl42_86), inference(avatar_component_clause, [], [f1026])).
fof(f1026, plain, (spl42_86 <=> (e21 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_86])])).
fof(f5766, plain, (~ (e21 = op2(e22, e23)) | (~ spl42_28 | ~ spl42_150 | spl42_234)), inference(forward_demodulation, [], [f5765, f433])).
fof(f433, plain, (e21 = h11(e13)), inference(cnf_transformation, [], [f26])).
fof(f5765, plain, (~ (op2(e22, e23) = h11(e13)) | (~ spl42_28 | ~ spl42_150 | spl42_234)), inference(forward_demodulation, [], [f5764, f749])).
fof(f749, plain, ((e13 = op1(e12, e11)) | ~ spl42_28), inference(avatar_component_clause, [], [f747])).
fof(f747, plain, (spl42_28 <=> (e13 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_28])])).
fof(f5764, plain, (~ (op2(e22, e23) = h11(op1(e12, e11))) | (~ spl42_150 | spl42_234)), inference(forward_demodulation, [], [f2019, f1436])).
fof(f1436, plain, ((e22 = h11(e12)) | ~ spl42_150), inference(avatar_component_clause, [], [f1435])).
fof(f1435, plain, (spl42_150 <=> (e22 = h11(e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_150])])).
fof(f2019, plain, (~ (h11(op1(e12, e11)) = op2(h11(e12), e23)) | spl42_234), inference(avatar_component_clause, [], [f2017])).
fof(f2017, plain, (spl42_234 <=> (h11(op1(e12, e11)) = op2(h11(e12), e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_234])])).
fof(f5708, plain, (~ spl42_46 | ~ spl42_84 | ~ spl42_148 | spl42_239), inference(avatar_contradiction_clause, [], [f5707])).
fof(f5707, plain, ($false | (~ spl42_46 | ~ spl42_84 | ~ spl42_148 | spl42_239)), inference(subsumption_resolution, [], [f5706, f1019])).
fof(f1019, plain, ((e23 = op2(e23, e20)) | ~ spl42_84), inference(avatar_component_clause, [], [f1017])).
fof(f1017, plain, (spl42_84 <=> (e23 = op2(e23, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_84])])).
fof(f5706, plain, (~ (e23 = op2(e23, e20)) | (~ spl42_46 | ~ spl42_148 | spl42_239)), inference(forward_demodulation, [], [f5705, f432])).
fof(f432, plain, (e23 = h11(e11)), inference(cnf_transformation, [], [f26])).
fof(f5705, plain, (~ (op2(e23, e20) = h11(e11)) | (~ spl42_46 | ~ spl42_148 | spl42_239)), inference(forward_demodulation, [], [f5704, f826])).
fof(f826, plain, ((e11 = op1(e11, e10)) | ~ spl42_46), inference(avatar_component_clause, [], [f824])).
fof(f824, plain, (spl42_46 <=> (e11 = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_46])])).
fof(f5704, plain, (~ (op2(e23, e20) = h11(op1(e11, e10))) | (~ spl42_148 | spl42_239)), inference(forward_demodulation, [], [f2039, f1426])).
fof(f2039, plain, (~ (h11(op1(e11, e10)) = op2(e23, h10(e10))) | spl42_239), inference(avatar_component_clause, [], [f2037])).
fof(f2037, plain, (spl42_239 <=> (h11(op1(e11, e10)) = op2(e23, h10(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl42_239])])).
fof(f5650, plain, (~ spl42_1 | ~ spl42_148 | ~ spl42_183 | spl42_229), inference(avatar_contradiction_clause, [], [f5649])).
fof(f5649, plain, ($false | (~ spl42_1 | ~ spl42_148 | ~ spl42_183 | spl42_229)), inference(subsumption_resolution, [], [f5648, f1640])).
fof(f1640, plain, ((e20 = h4(e10)) | ~ spl42_183), inference(avatar_component_clause, [], [f1639])).
fof(f1639, plain, (spl42_183 <=> (e20 = h4(e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_183])])).
fof(f5648, plain, (~ (e20 = h4(e10)) | (~ spl42_1 | ~ spl42_148 | spl42_229)), inference(forward_demodulation, [], [f5647, f4429])).
fof(f5647, plain, (~ (h4(e10) = h11(e10)) | (~ spl42_1 | spl42_229)), inference(forward_demodulation, [], [f1999, f635])).
fof(f635, plain, ((e10 = op1(e13, e13)) | ~ spl42_1), inference(avatar_component_clause, [], [f633])).
fof(f633, plain, (spl42_1 <=> (e10 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_1])])).
fof(f1999, plain, (~ (h4(e10) = h11(op1(e13, e13))) | spl42_229), inference(avatar_component_clause, [], [f1997])).
fof(f1997, plain, (spl42_229 <=> (h4(e10) = h11(op1(e13, e13)))), introduced(avatar_definition, [new_symbols(naming, [spl42_229])])).
fof(f5588, plain, (~ spl42_35 | ~ spl42_150 | ~ spl42_176 | spl42_236), inference(avatar_contradiction_clause, [], [f5587])).
fof(f5587, plain, ($false | (~ spl42_35 | ~ spl42_150 | ~ spl42_176 | spl42_236)), inference(subsumption_resolution, [], [f5586, f1598])).
fof(f1598, plain, ((e22 = h6(e12)) | ~ spl42_176), inference(avatar_component_clause, [], [f1597])).
fof(f1597, plain, (spl42_176 <=> (e22 = h6(e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_176])])).
fof(f5586, plain, (~ (e22 = h6(e12)) | (~ spl42_35 | ~ spl42_150 | spl42_236)), inference(forward_demodulation, [], [f5585, f1436])).
fof(f5585, plain, (~ (h6(e12) = h11(e12)) | (~ spl42_35 | spl42_236)), inference(forward_demodulation, [], [f2027, f779])).
fof(f779, plain, ((e12 = op1(e11, e13)) | ~ spl42_35), inference(avatar_component_clause, [], [f777])).
fof(f777, plain, (spl42_35 <=> (e12 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_35])])).
fof(f2027, plain, (~ (h6(e12) = h11(op1(e11, e13))) | spl42_236), inference(avatar_component_clause, [], [f2025])).
fof(f2025, plain, (spl42_236 <=> (h6(e12) = h11(op1(e11, e13)))), introduced(avatar_definition, [new_symbols(naming, [spl42_236])])).
fof(f5527, plain, (~ spl42_40 | ~ spl42_74 | ~ spl42_150 | spl42_237), inference(avatar_contradiction_clause, [], [f5526])).
fof(f5526, plain, ($false | (~ spl42_40 | ~ spl42_74 | ~ spl42_150 | spl42_237)), inference(subsumption_resolution, [], [f5525, f977])).
fof(f977, plain, ((e21 = op2(e23, e22)) | ~ spl42_74), inference(avatar_component_clause, [], [f975])).
fof(f975, plain, (spl42_74 <=> (e21 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_74])])).
fof(f5525, plain, (~ (e21 = op2(e23, e22)) | (~ spl42_40 | ~ spl42_150 | spl42_237)), inference(forward_demodulation, [], [f5524, f433])).
fof(f5524, plain, (~ (op2(e23, e22) = h11(e13)) | (~ spl42_40 | ~ spl42_150 | spl42_237)), inference(forward_demodulation, [], [f5523, f800])).
fof(f800, plain, ((e13 = op1(e11, e12)) | ~ spl42_40), inference(avatar_component_clause, [], [f798])).
fof(f798, plain, (spl42_40 <=> (e13 = op1(e11, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_40])])).
fof(f5523, plain, (~ (op2(e23, e22) = h11(op1(e11, e12))) | (~ spl42_150 | spl42_237)), inference(forward_demodulation, [], [f2031, f1436])).
fof(f2031, plain, (~ (h11(op1(e11, e12)) = op2(e23, h11(e12))) | spl42_237), inference(avatar_component_clause, [], [f2029])).
fof(f2029, plain, (spl42_237 <=> (h11(op1(e11, e12)) = op2(e23, h11(e12)))), introduced(avatar_definition, [new_symbols(naming, [spl42_237])])).
fof(f5460, plain, (~ spl42_58 | ~ spl42_120 | ~ spl42_148 | spl42_242), inference(avatar_contradiction_clause, [], [f5459])).
fof(f5459, plain, ($false | (~ spl42_58 | ~ spl42_120 | ~ spl42_148 | spl42_242)), inference(subsumption_resolution, [], [f5458, f1172])).
fof(f1172, plain, ((e23 = op2(e20, e23)) | ~ spl42_120), inference(avatar_component_clause, [], [f1170])).
fof(f1170, plain, (spl42_120 <=> (e23 = op2(e20, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_120])])).
fof(f5458, plain, (~ (e23 = op2(e20, e23)) | (~ spl42_58 | ~ spl42_148 | spl42_242)), inference(forward_demodulation, [], [f5457, f432])).
fof(f5457, plain, (~ (op2(e20, e23) = h11(e11)) | (~ spl42_58 | ~ spl42_148 | spl42_242)), inference(forward_demodulation, [], [f5456, f877])).
fof(f877, plain, ((e11 = op1(e10, e11)) | ~ spl42_58), inference(avatar_component_clause, [], [f875])).
fof(f875, plain, (spl42_58 <=> (e11 = op1(e10, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_58])])).
fof(f5456, plain, (~ (op2(e20, e23) = h11(op1(e10, e11))) | (~ spl42_148 | spl42_242)), inference(forward_demodulation, [], [f2051, f1426])).
fof(f2051, plain, (~ (h11(op1(e10, e11)) = op2(h10(e10), e23)) | spl42_242), inference(avatar_component_clause, [], [f2049])).
fof(f2049, plain, (spl42_242 <=> (h11(op1(e10, e11)) = op2(h10(e10), e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_242])])).
fof(f5402, plain, (~ spl42_52 | ~ spl42_126 | ~ spl42_148 | spl42_240), inference(avatar_contradiction_clause, [], [f5401])).
fof(f5401, plain, ($false | (~ spl42_52 | ~ spl42_126 | ~ spl42_148 | spl42_240)), inference(subsumption_resolution, [], [f5400, f1198])).
fof(f1198, plain, ((e21 = op2(e20, e21)) | ~ spl42_126), inference(avatar_component_clause, [], [f1196])).
fof(f1196, plain, (spl42_126 <=> (e21 = op2(e20, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_126])])).
fof(f5400, plain, (~ (e21 = op2(e20, e21)) | (~ spl42_52 | ~ spl42_148 | spl42_240)), inference(forward_demodulation, [], [f5399, f433])).
fof(f5399, plain, (~ (op2(e20, e21) = h11(e13)) | (~ spl42_52 | ~ spl42_148 | spl42_240)), inference(forward_demodulation, [], [f5398, f851])).
fof(f851, plain, ((e13 = op1(e10, e13)) | ~ spl42_52), inference(avatar_component_clause, [], [f849])).
fof(f849, plain, (spl42_52 <=> (e13 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_52])])).
fof(f5398, plain, (~ (op2(e20, e21) = h11(op1(e10, e13))) | (~ spl42_148 | spl42_240)), inference(forward_demodulation, [], [f2043, f1426])).
fof(f2043, plain, (~ (h11(op1(e10, e13)) = op2(h10(e10), e21)) | spl42_240), inference(avatar_component_clause, [], [f2041])).
fof(f2041, plain, (spl42_240 <=> (h11(op1(e10, e13)) = op2(h10(e10), e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_240])])).
fof(f5346, plain, (~ spl42_55 | ~ spl42_123 | ~ spl42_148 | ~ spl42_150 | spl42_241), inference(avatar_contradiction_clause, [], [f5345])).
fof(f5345, plain, ($false | (~ spl42_55 | ~ spl42_123 | ~ spl42_148 | ~ spl42_150 | spl42_241)), inference(subsumption_resolution, [], [f5344, f1185])).
fof(f1185, plain, ((e22 = op2(e20, e22)) | ~ spl42_123), inference(avatar_component_clause, [], [f1183])).
fof(f1183, plain, (spl42_123 <=> (e22 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_123])])).
fof(f5344, plain, (~ (e22 = op2(e20, e22)) | (~ spl42_55 | ~ spl42_148 | ~ spl42_150 | spl42_241)), inference(forward_demodulation, [], [f5343, f1436])).
fof(f5343, plain, (~ (op2(e20, e22) = h11(e12)) | (~ spl42_55 | ~ spl42_148 | ~ spl42_150 | spl42_241)), inference(forward_demodulation, [], [f5342, f864])).
fof(f864, plain, ((e12 = op1(e10, e12)) | ~ spl42_55), inference(avatar_component_clause, [], [f862])).
fof(f862, plain, (spl42_55 <=> (e12 = op1(e10, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_55])])).
fof(f5342, plain, (~ (op2(e20, e22) = h11(op1(e10, e12))) | (~ spl42_148 | ~ spl42_150 | spl42_241)), inference(forward_demodulation, [], [f5341, f1426])).
fof(f5341, plain, (~ (h11(op1(e10, e12)) = op2(h10(e10), e22)) | (~ spl42_150 | spl42_241)), inference(forward_demodulation, [], [f2047, f1436])).
fof(f2047, plain, (~ (h11(op1(e10, e12)) = op2(h10(e10), h11(e12))) | spl42_241), inference(avatar_component_clause, [], [f2045])).
fof(f2045, plain, (spl42_241 <=> (h11(op1(e10, e12)) = op2(h10(e10), h11(e12)))), introduced(avatar_definition, [new_symbols(naming, [spl42_241])])).
fof(f5288, plain, (~ spl42_31 | ~ spl42_99 | ~ spl42_148 | ~ spl42_150 | spl42_235), inference(avatar_contradiction_clause, [], [f5287])).
fof(f5287, plain, ($false | (~ spl42_31 | ~ spl42_99 | ~ spl42_148 | ~ spl42_150 | spl42_235)), inference(subsumption_resolution, [], [f5286, f1083])).
fof(f1083, plain, ((e22 = op2(e22, e20)) | ~ spl42_99), inference(avatar_component_clause, [], [f1081])).
fof(f1081, plain, (spl42_99 <=> (e22 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_99])])).
fof(f5286, plain, (~ (e22 = op2(e22, e20)) | (~ spl42_31 | ~ spl42_148 | ~ spl42_150 | spl42_235)), inference(forward_demodulation, [], [f5285, f1436])).
fof(f5285, plain, (~ (op2(e22, e20) = h11(e12)) | (~ spl42_31 | ~ spl42_148 | ~ spl42_150 | spl42_235)), inference(forward_demodulation, [], [f5284, f762])).
fof(f762, plain, ((e12 = op1(e12, e10)) | ~ spl42_31), inference(avatar_component_clause, [], [f760])).
fof(f760, plain, (spl42_31 <=> (e12 = op1(e12, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_31])])).
fof(f5284, plain, (~ (op2(e22, e20) = h11(op1(e12, e10))) | (~ spl42_148 | ~ spl42_150 | spl42_235)), inference(forward_demodulation, [], [f5283, f1436])).
fof(f5283, plain, (~ (h11(op1(e12, e10)) = op2(h11(e12), e20)) | (~ spl42_148 | spl42_235)), inference(forward_demodulation, [], [f2023, f1426])).
fof(f2023, plain, (~ (h11(op1(e12, e10)) = op2(h11(e12), h10(e10))) | spl42_235), inference(avatar_component_clause, [], [f2021])).
fof(f2021, plain, (spl42_235 <=> (h11(op1(e12, e10)) = op2(h11(e12), h10(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl42_235])])).
fof(f5226, plain, (~ spl42_6 | ~ spl42_108 | ~ spl42_150 | spl42_230), inference(avatar_contradiction_clause, [], [f5225])).
fof(f5225, plain, ($false | (~ spl42_6 | ~ spl42_108 | ~ spl42_150 | spl42_230)), inference(subsumption_resolution, [], [f5224, f1121])).
fof(f1121, plain, ((e23 = op2(e21, e22)) | ~ spl42_108), inference(avatar_component_clause, [], [f1119])).
fof(f1119, plain, (spl42_108 <=> (e23 = op2(e21, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_108])])).
fof(f5224, plain, (~ (e23 = op2(e21, e22)) | (~ spl42_6 | ~ spl42_150 | spl42_230)), inference(forward_demodulation, [], [f5223, f432])).
fof(f5223, plain, (~ (op2(e21, e22) = h11(e11)) | (~ spl42_6 | ~ spl42_150 | spl42_230)), inference(forward_demodulation, [], [f5222, f656])).
fof(f656, plain, ((e11 = op1(e13, e12)) | ~ spl42_6), inference(avatar_component_clause, [], [f654])).
fof(f654, plain, (spl42_6 <=> (e11 = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_6])])).
fof(f5222, plain, (~ (op2(e21, e22) = h11(op1(e13, e12))) | (~ spl42_150 | spl42_230)), inference(forward_demodulation, [], [f2003, f1436])).
fof(f2003, plain, (~ (h11(op1(e13, e12)) = op2(e21, h11(e12))) | spl42_230), inference(avatar_component_clause, [], [f2001])).
fof(f2001, plain, (spl42_230 <=> (h11(op1(e13, e12)) = op2(e21, h11(e12)))), introduced(avatar_definition, [new_symbols(naming, [spl42_230])])).
fof(f5175, plain, (~ spl42_21 | ~ spl42_89 | ~ spl42_148 | ~ spl42_150 | spl42_227), inference(avatar_contradiction_clause, [], [f5174])).
fof(f5174, plain, ($false | (~ spl42_21 | ~ spl42_89 | ~ spl42_148 | ~ spl42_150 | spl42_227)), inference(subsumption_resolution, [], [f5170, f1041])).
fof(f1041, plain, ((e20 = op2(e22, e22)) | ~ spl42_89), inference(avatar_component_clause, [], [f1039])).
fof(f1039, plain, (spl42_89 <=> (e20 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_89])])).
fof(f5170, plain, (~ (e20 = op2(e22, e22)) | (~ spl42_21 | ~ spl42_148 | ~ spl42_150 | spl42_227)), inference(backward_demodulation, [], [f5161, f1436])).
fof(f5161, plain, (~ (e20 = op2(h11(e12), h11(e12))) | (~ spl42_21 | ~ spl42_148 | spl42_227)), inference(forward_demodulation, [], [f5160, f4429])).
fof(f5160, plain, (~ (h11(e10) = op2(h11(e12), h11(e12))) | (~ spl42_21 | spl42_227)), inference(forward_demodulation, [], [f1991, f720])).
fof(f720, plain, ((e10 = op1(e12, e12)) | ~ spl42_21), inference(avatar_component_clause, [], [f718])).
fof(f718, plain, (spl42_21 <=> (e10 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_21])])).
fof(f1991, plain, (~ (h11(op1(e12, e12)) = op2(h11(e12), h11(e12))) | spl42_227), inference(avatar_component_clause, [], [f1989])).
fof(f1989, plain, (spl42_227 <=> (h11(op1(e12, e12)) = op2(h11(e12), h11(e12)))), introduced(avatar_definition, [new_symbols(naming, [spl42_227])])).
fof(f5102, plain, (~ spl42_119 | ~ spl42_120), inference(avatar_contradiction_clause, [], [f5101])).
fof(f5101, plain, ($false | (~ spl42_119 | ~ spl42_120)), inference(subsumption_resolution, [], [f5100, f339])).
fof(f339, plain, ~ (e22 = e23), inference(cnf_transformation, [], [f10])).
fof(f10, plain, (~ (e22 = e23) & ~ (e21 = e23) & ~ (e21 = e22) & ~ (e20 = e23) & ~ (e20 = e22) & ~ (e20 = e21)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG037+1.p', ax10)).
fof(f5100, plain, ((e22 = e23) | (~ spl42_119 | ~ spl42_120)), inference(backward_demodulation, [], [f1172, f1168])).
fof(f1168, plain, ((e22 = op2(e20, e23)) | ~ spl42_119), inference(avatar_component_clause, [], [f1166])).
fof(f1166, plain, (spl42_119 <=> (e22 = op2(e20, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_119])])).
fof(f5046, plain, (~ spl42_102 | ~ spl42_114), inference(avatar_split_clause, [], [f5045, f1145, f1094])).
fof(f1094, plain, (spl42_102 <=> (e21 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_102])])).
fof(f1145, plain, (spl42_114 <=> (e21 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_114])])).
fof(f5045, plain, (~ (e21 = op2(e21, e23)) | ~ spl42_114), inference(forward_demodulation, [], [f1339, f4880])).
fof(f4880, plain, ((e21 = h1(e12)) | ~ spl42_114), inference(backward_demodulation, [], [f395, f1147])).
fof(f1147, plain, ((e21 = op2(e21, e20)) | ~ spl42_114), inference(avatar_component_clause, [], [f1145])).
fof(f395, plain, (op2(e21, e20) = h1(e12)), inference(cnf_transformation, [], [f16])).
fof(f16, plain, ((op2(e21, e20) = h1(e12)) & (op2(e20, e20) = h1(e10)) & (e21 = h1(e13)) & (e20 = h1(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG037+1.p', ax16)).
fof(f1339, plain, ~ (op2(e21, e23) = h1(e12)), inference(backward_demodulation, [], [f312, f395])).
fof(f312, plain, ~ (op2(e21, e20) = op2(e21, e23)), inference(cnf_transformation, [], [f8])).
fof(f8, plain, (~ (op2(e23, e22) = op2(e23, e23)) & ~ (op2(e23, e21) = op2(e23, e23)) & ~ (op2(e23, e21) = op2(e23, e22)) & ~ (op2(e23, e20) = op2(e23, e23)) & ~ (op2(e23, e20) = op2(e23, e22)) & ~ (op2(e23, e20) = op2(e23, e21)) & ~ (op2(e22, e22) = op2(e22, e23)) & ~ (op2(e22, e21) = op2(e22, e23)) & ~ (op2(e22, e21) = op2(e22, e22)) & ~ (op2(e22, e20) = op2(e22, e23)) & ~ (op2(e22, e20) = op2(e22, e22)) & ~ (op2(e22, e20) = op2(e22, e21)) & ~ (op2(e21, e22) = op2(e21, e23)) & ~ (op2(e21, e21) = op2(e21, e23)) & ~ (op2(e21, e21) = op2(e21, e22)) & ~ (op2(e21, e20) = op2(e21, e23)) & ~ (op2(e21, e20) = op2(e21, e22)) & ~ (op2(e21, e20) = op2(e21, e21)) & ~ (op2(e20, e22) = op2(e20, e23)) & ~ (op2(e20, e21) = op2(e20, e23)) & ~ (op2(e20, e21) = op2(e20, e22)) & ~ (op2(e20, e20) = op2(e20, e23)) & ~ (op2(e20, e20) = op2(e20, e22)) & ~ (op2(e20, e20) = op2(e20, e21)) & ~ (op2(e22, e23) = op2(e23, e23)) & ~ (op2(e21, e23) = op2(e23, e23)) & ~ (op2(e21, e23) = op2(e22, e23)) & ~ (op2(e20, e23) = op2(e23, e23)) & ~ (op2(e20, e23) = op2(e22, e23)) & ~ (op2(e20, e23) = op2(e21, e23)) & ~ (op2(e22, e22) = op2(e23, e22)) & ~ (op2(e21, e22) = op2(e23, e22)) & ~ (op2(e21, e22) = op2(e22, e22)) & ~ (op2(e20, e22) = op2(e23, e22)) & ~ (op2(e20, e22) = op2(e22, e22)) & ~ (op2(e20, e22) = op2(e21, e22)) & ~ (op2(e22, e21) = op2(e23, e21)) & ~ (op2(e21, e21) = op2(e23, e21)) & ~ (op2(e21, e21) = op2(e22, e21)) & ~ (op2(e20, e21) = op2(e23, e21)) & ~ (op2(e20, e21) = op2(e22, e21)) & ~ (op2(e20, e21) = op2(e21, e21)) & ~ (op2(e22, e20) = op2(e23, e20)) & ~ (op2(e21, e20) = op2(e23, e20)) & ~ (op2(e21, e20) = op2(e22, e20)) & ~ (op2(e20, e20) = op2(e23, e20)) & ~ (op2(e20, e20) = op2(e22, e20)) & ~ (op2(e20, e20) = op2(e21, e20))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG037+1.p', ax8)).
fof(f5042, plain, (~ spl42_104 | ~ spl42_248), inference(avatar_split_clause, [], [f5041, f2151, f1102])).
fof(f1102, plain, (spl42_104 <=> (e23 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_104])])).
fof(f2151, plain, (spl42_248 <=> (e23 = h10(e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_248])])).
fof(f5041, plain, (~ (e23 = op2(e21, e23)) | ~ spl42_248), inference(forward_demodulation, [], [f1385, f2152])).
fof(f2152, plain, ((e23 = h10(e12)) | ~ spl42_248), inference(avatar_component_clause, [], [f2151])).
fof(f1385, plain, ~ (op2(e21, e23) = h10(e12)), inference(backward_demodulation, [], [f298, f431])).
fof(f431, plain, (op2(e20, e23) = h10(e12)), inference(cnf_transformation, [], [f25])).
fof(f298, plain, ~ (op2(e20, e23) = op2(e21, e23)), inference(cnf_transformation, [], [f8])).
fof(f5039, plain, (~ spl42_18 | ~ spl42_96 | ~ spl42_150 | spl42_233), inference(avatar_contradiction_clause, [], [f5038])).
fof(f5038, plain, ($false | (~ spl42_18 | ~ spl42_96 | ~ spl42_150 | spl42_233)), inference(subsumption_resolution, [], [f5037, f1070])).
fof(f1070, plain, ((e23 = op2(e22, e21)) | ~ spl42_96), inference(avatar_component_clause, [], [f1068])).
fof(f1068, plain, (spl42_96 <=> (e23 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_96])])).
fof(f5037, plain, (~ (e23 = op2(e22, e21)) | (~ spl42_18 | ~ spl42_150 | spl42_233)), inference(forward_demodulation, [], [f5036, f432])).
fof(f5036, plain, (~ (op2(e22, e21) = h11(e11)) | (~ spl42_18 | ~ spl42_150 | spl42_233)), inference(forward_demodulation, [], [f5035, f707])).
fof(f707, plain, ((e11 = op1(e12, e13)) | ~ spl42_18), inference(avatar_component_clause, [], [f705])).
fof(f705, plain, (spl42_18 <=> (e11 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_18])])).
fof(f5035, plain, (~ (op2(e22, e21) = h11(op1(e12, e13))) | (~ spl42_150 | spl42_233)), inference(forward_demodulation, [], [f2015, f1436])).
fof(f2015, plain, (~ (h11(op1(e12, e13)) = op2(h11(e12), e21)) | spl42_233), inference(avatar_component_clause, [], [f2013])).
fof(f2013, plain, (spl42_233 <=> (h11(op1(e12, e13)) = op2(h11(e12), e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_233])])).
fof(f4947, plain, (~ spl42_19 | ~ spl42_31), inference(avatar_split_clause, [], [f4944, f760, f709])).
fof(f709, plain, (spl42_19 <=> (e12 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_19])])).
fof(f4944, plain, (~ (e12 = op1(e12, e13)) | ~ spl42_31), inference(backward_demodulation, [], [f270, f762])).
fof(f270, plain, ~ (op1(e12, e10) = op1(e12, e13)), inference(cnf_transformation, [], [f7])).
fof(f7, plain, (~ (op1(e13, e12) = op1(e13, e13)) & ~ (op1(e13, e11) = op1(e13, e13)) & ~ (op1(e13, e11) = op1(e13, e12)) & ~ (op1(e13, e10) = op1(e13, e13)) & ~ (op1(e13, e10) = op1(e13, e12)) & ~ (op1(e13, e10) = op1(e13, e11)) & ~ (op1(e12, e12) = op1(e12, e13)) & ~ (op1(e12, e11) = op1(e12, e13)) & ~ (op1(e12, e11) = op1(e12, e12)) & ~ (op1(e12, e10) = op1(e12, e13)) & ~ (op1(e12, e10) = op1(e12, e12)) & ~ (op1(e12, e10) = op1(e12, e11)) & ~ (op1(e11, e12) = op1(e11, e13)) & ~ (op1(e11, e11) = op1(e11, e13)) & ~ (op1(e11, e11) = op1(e11, e12)) & ~ (op1(e11, e10) = op1(e11, e13)) & ~ (op1(e11, e10) = op1(e11, e12)) & ~ (op1(e11, e10) = op1(e11, e11)) & ~ (op1(e10, e12) = op1(e10, e13)) & ~ (op1(e10, e11) = op1(e10, e13)) & ~ (op1(e10, e11) = op1(e10, e12)) & ~ (op1(e10, e10) = op1(e10, e13)) & ~ (op1(e10, e10) = op1(e10, e12)) & ~ (op1(e10, e10) = op1(e10, e11)) & ~ (op1(e12, e13) = op1(e13, e13)) & ~ (op1(e11, e13) = op1(e13, e13)) & ~ (op1(e11, e13) = op1(e12, e13)) & ~ (op1(e10, e13) = op1(e13, e13)) & ~ (op1(e10, e13) = op1(e12, e13)) & ~ (op1(e10, e13) = op1(e11, e13)) & ~ (op1(e12, e12) = op1(e13, e12)) & ~ (op1(e11, e12) = op1(e13, e12)) & ~ (op1(e11, e12) = op1(e12, e12)) & ~ (op1(e10, e12) = op1(e13, e12)) & ~ (op1(e10, e12) = op1(e12, e12)) & ~ (op1(e10, e12) = op1(e11, e12)) & ~ (op1(e12, e11) = op1(e13, e11)) & ~ (op1(e11, e11) = op1(e13, e11)) & ~ (op1(e11, e11) = op1(e12, e11)) & ~ (op1(e10, e11) = op1(e13, e11)) & ~ (op1(e10, e11) = op1(e12, e11)) & ~ (op1(e10, e11) = op1(e11, e11)) & ~ (op1(e12, e10) = op1(e13, e10)) & ~ (op1(e11, e10) = op1(e13, e10)) & ~ (op1(e11, e10) = op1(e12, e10)) & ~ (op1(e10, e10) = op1(e13, e10)) & ~ (op1(e10, e10) = op1(e12, e10)) & ~ (op1(e10, e10) = op1(e11, e10))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG037+1.p', ax7)).
fof(f4934, plain, (~ spl42_38 | ~ spl42_46), inference(avatar_split_clause, [], [f4929, f824, f790])).
fof(f790, plain, (spl42_38 <=> (e11 = op1(e11, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_38])])).
fof(f4929, plain, (~ (e11 = op1(e11, e12)) | ~ spl42_46), inference(backward_demodulation, [], [f263, f826])).
fof(f263, plain, ~ (op1(e11, e10) = op1(e11, e12)), inference(cnf_transformation, [], [f7])).
fof(f4916, plain, (spl42_16 | ~ spl42_65), inference(avatar_split_clause, [], [f4906, f905, f696])).
fof(f696, plain, (spl42_16 <=> (e13 = op1(e13, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_16])])).
fof(f905, plain, (spl42_65 <=> (e10 = unit1)), introduced(avatar_definition, [new_symbols(naming, [spl42_65])])).
fof(f4906, plain, ((e13 = op1(e13, e10)) | ~ spl42_65), inference(backward_demodulation, [], [f141, f907])).
fof(f907, plain, ((e10 = unit1) | ~ spl42_65), inference(avatar_component_clause, [], [f905])).
fof(f141, plain, (e13 = op1(e13, unit1)), inference(cnf_transformation, [], [f2])).
fof(f2, plain, (((e13 = unit1) | (e12 = unit1) | (e11 = unit1) | (e10 = unit1)) & (e13 = op1(e13, unit1)) & (e13 = op1(unit1, e13)) & (e12 = op1(e12, unit1)) & (e12 = op1(unit1, e12)) & (e11 = op1(e11, unit1)) & (e11 = op1(unit1, e11)) & (e10 = op1(e10, unit1)) & (e10 = op1(unit1, e10))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG037+1.p', ax2)).
fof(f4915, plain, (~ spl42_50 | ~ spl42_65), inference(avatar_contradiction_clause, [], [f4914])).
fof(f4914, plain, ($false | (~ spl42_50 | ~ spl42_65)), inference(subsumption_resolution, [], [f4913, f332])).
fof(f332, plain, ~ (e11 = e13), inference(cnf_transformation, [], [f9])).
fof(f9, plain, (~ (e12 = e13) & ~ (e11 = e13) & ~ (e11 = e12) & ~ (e10 = e13) & ~ (e10 = e12) & ~ (e10 = e11)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG037+1.p', ax9)).
fof(f4913, plain, ((e11 = e13) | (~ spl42_50 | ~ spl42_65)), inference(forward_demodulation, [], [f4905, f843])).
fof(f843, plain, ((e11 = op1(e10, e13)) | ~ spl42_50), inference(avatar_component_clause, [], [f841])).
fof(f841, plain, (spl42_50 <=> (e11 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_50])])).
fof(f4905, plain, ((e13 = op1(e10, e13)) | ~ spl42_65), inference(backward_demodulation, [], [f140, f907])).
fof(f140, plain, (e13 = op1(unit1, e13)), inference(cnf_transformation, [], [f2])).
fof(f4912, plain, (spl42_31 | ~ spl42_65), inference(avatar_split_clause, [], [f4904, f905, f760])).
fof(f4904, plain, ((e12 = op1(e12, e10)) | ~ spl42_65), inference(backward_demodulation, [], [f139, f907])).
fof(f139, plain, (e12 = op1(e12, unit1)), inference(cnf_transformation, [], [f2])).
fof(f4911, plain, (spl42_46 | ~ spl42_65), inference(avatar_split_clause, [], [f4902, f905, f824])).
fof(f4902, plain, ((e11 = op1(e11, e10)) | ~ spl42_65), inference(backward_demodulation, [], [f137, f907])).
fof(f137, plain, (e11 = op1(e11, unit1)), inference(cnf_transformation, [], [f2])).
fof(f4910, plain, (spl42_58 | ~ spl42_65), inference(avatar_split_clause, [], [f4901, f905, f875])).
fof(f4901, plain, ((e11 = op1(e10, e11)) | ~ spl42_65), inference(backward_demodulation, [], [f136, f907])).
fof(f136, plain, (e11 = op1(unit1, e11)), inference(cnf_transformation, [], [f2])).
fof(f4876, plain, (spl42_84 | ~ spl42_133), inference(avatar_split_clause, [], [f4869, f1226, f1017])).
fof(f1226, plain, (spl42_133 <=> (e20 = unit2)), introduced(avatar_definition, [new_symbols(naming, [spl42_133])])).
fof(f4869, plain, ((e23 = op2(e23, e20)) | ~ spl42_133), inference(backward_demodulation, [], [f198, f1228])).
fof(f1228, plain, ((e20 = unit2) | ~ spl42_133), inference(avatar_component_clause, [], [f1226])).
fof(f198, plain, (e23 = op2(e23, unit2)), inference(cnf_transformation, [], [f5])).
fof(f5, plain, (((e23 = unit2) | (e22 = unit2) | (e21 = unit2) | (e20 = unit2)) & (e23 = op2(e23, unit2)) & (e23 = op2(unit2, e23)) & (e22 = op2(e22, unit2)) & (e22 = op2(unit2, e22)) & (e21 = op2(e21, unit2)) & (e21 = op2(unit2, e21)) & (e20 = op2(e20, unit2)) & (e20 = op2(unit2, e20))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG037+1.p', ax5)).
fof(f4875, plain, (spl42_120 | ~ spl42_133), inference(avatar_split_clause, [], [f4868, f1226, f1170])).
fof(f4868, plain, ((e23 = op2(e20, e23)) | ~ spl42_133), inference(backward_demodulation, [], [f197, f1228])).
fof(f197, plain, (e23 = op2(unit2, e23)), inference(cnf_transformation, [], [f5])).
fof(f4874, plain, (spl42_114 | ~ spl42_133), inference(avatar_split_clause, [], [f4865, f1226, f1145])).
fof(f4865, plain, ((e21 = op2(e21, e20)) | ~ spl42_133), inference(backward_demodulation, [], [f194, f1228])).
fof(f194, plain, (e21 = op2(e21, unit2)), inference(cnf_transformation, [], [f5])).
fof(f4873, plain, (spl42_126 | ~ spl42_133), inference(avatar_split_clause, [], [f4864, f1226, f1196])).
fof(f4864, plain, ((e21 = op2(e20, e21)) | ~ spl42_133), inference(backward_demodulation, [], [f193, f1228])).
fof(f193, plain, (e21 = op2(unit2, e21)), inference(cnf_transformation, [], [f5])).
fof(f4837, plain, (~ spl42_11 | spl42_231), inference(avatar_contradiction_clause, [], [f4836])).
fof(f4836, plain, ($false | (~ spl42_11 | spl42_231)), inference(trivial_inequality_removal, [], [f4835])).
fof(f4835, plain, (~ (h11(e12) = h11(e12)) | (~ spl42_11 | spl42_231)), inference(forward_demodulation, [], [f2007, f677])).
fof(f677, plain, ((e12 = op1(e13, e11)) | ~ spl42_11), inference(avatar_component_clause, [], [f675])).
fof(f675, plain, (spl42_11 <=> (e12 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_11])])).
fof(f2007, plain, (~ (h11(e12) = h11(op1(e13, e11))) | spl42_231), inference(avatar_component_clause, [], [f2005])).
fof(f2005, plain, (spl42_231 <=> (h11(e12) = h11(op1(e13, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl42_231])])).
fof(f4766, plain, (~ spl42_39 | ~ spl42_55), inference(avatar_split_clause, [], [f4765, f862, f794])).
fof(f794, plain, (spl42_39 <=> (e12 = op1(e11, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_39])])).
fof(f4765, plain, (~ (e12 = op1(e11, e12)) | ~ spl42_55), inference(forward_demodulation, [], [f244, f864])).
fof(f244, plain, ~ (op1(e10, e12) = op1(e11, e12)), inference(cnf_transformation, [], [f7])).
fof(f4706, plain, (spl42_10 | ~ spl42_68), inference(avatar_contradiction_clause, [], [f4705])).
fof(f4705, plain, ($false | (spl42_10 | ~ spl42_68)), inference(subsumption_resolution, [], [f4695, f672])).
fof(f672, plain, (~ (e11 = op1(e13, e11)) | spl42_10), inference(avatar_component_clause, [], [f671])).
fof(f671, plain, (spl42_10 <=> (e11 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_10])])).
fof(f4695, plain, ((e11 = op1(e13, e11)) | ~ spl42_68), inference(backward_demodulation, [], [f136, f919])).
fof(f919, plain, ((e13 = unit1) | ~ spl42_68), inference(avatar_component_clause, [], [f917])).
fof(f917, plain, (spl42_68 <=> (e13 = unit1)), introduced(avatar_definition, [new_symbols(naming, [spl42_68])])).
fof(f4680, plain, (spl42_404 | ~ spl42_116), inference(avatar_split_clause, [], [f4679, f1153, f3346])).
fof(f3346, plain, (spl42_404 <=> (e23 = h1(e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_404])])).
fof(f1153, plain, (spl42_116 <=> (e23 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_116])])).
fof(f4679, plain, ((e23 = h1(e12)) | ~ spl42_116), inference(backward_demodulation, [], [f395, f1155])).
fof(f1155, plain, ((e23 = op2(e21, e20)) | ~ spl42_116), inference(avatar_component_clause, [], [f1153])).
fof(f4664, plain, (spl42_78 | ~ spl42_136), inference(avatar_contradiction_clause, [], [f4663])).
fof(f4663, plain, ($false | (spl42_78 | ~ spl42_136)), inference(subsumption_resolution, [], [f4654, f993])).
fof(f993, plain, (~ (e21 = op2(e23, e21)) | spl42_78), inference(avatar_component_clause, [], [f992])).
fof(f992, plain, (spl42_78 <=> (e21 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_78])])).
fof(f4654, plain, ((e21 = op2(e23, e21)) | ~ spl42_136), inference(backward_demodulation, [], [f193, f1240])).
fof(f1240, plain, ((e23 = unit2) | ~ spl42_136), inference(avatar_component_clause, [], [f1238])).
fof(f1238, plain, (spl42_136 <=> (e23 = unit2)), introduced(avatar_definition, [new_symbols(naming, [spl42_136])])).
fof(f4606, plain, (spl42_352 | ~ spl42_128), inference(avatar_split_clause, [], [f4605, f1204, f2954])).
fof(f2954, plain, (spl42_352 <=> (e23 = h4(e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_352])])).
fof(f1204, plain, (spl42_128 <=> (e23 = op2(e20, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_128])])).
fof(f4605, plain, ((e23 = h4(e12)) | ~ spl42_128), inference(forward_demodulation, [], [f407, f1206])).
fof(f1206, plain, ((e23 = op2(e20, e21)) | ~ spl42_128), inference(avatar_component_clause, [], [f1204])).
fof(f407, plain, (op2(e20, e21) = h4(e12)), inference(cnf_transformation, [], [f19])).
fof(f19, plain, ((op2(e20, e21) = h4(e12)) & (op2(e21, e21) = h4(e10)) & (e20 = h4(e13)) & (e21 = h4(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG037+1.p', ax19)).
fof(f4547, plain, (~ spl42_1 | ~ spl42_2), inference(avatar_contradiction_clause, [], [f4546])).
fof(f4546, plain, ($false | (~ spl42_1 | ~ spl42_2)), inference(subsumption_resolution, [], [f4545, f328])).
fof(f328, plain, ~ (e10 = e11), inference(cnf_transformation, [], [f9])).
fof(f4545, plain, ((e10 = e11) | (~ spl42_1 | ~ spl42_2)), inference(forward_demodulation, [], [f639, f635])).
fof(f639, plain, ((e11 = op1(e13, e13)) | ~ spl42_2), inference(avatar_component_clause, [], [f637])).
fof(f637, plain, (spl42_2 <=> (e11 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_2])])).
fof(f4525, plain, (~ spl42_41 | ~ spl42_42), inference(avatar_contradiction_clause, [], [f4524])).
fof(f4524, plain, ($false | (~ spl42_41 | ~ spl42_42)), inference(subsumption_resolution, [], [f4523, f328])).
fof(f4523, plain, ((e10 = e11) | (~ spl42_41 | ~ spl42_42)), inference(forward_demodulation, [], [f809, f805])).
fof(f805, plain, ((e10 = op1(e11, e11)) | ~ spl42_41), inference(avatar_component_clause, [], [f803])).
fof(f803, plain, (spl42_41 <=> (e10 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_41])])).
fof(f809, plain, ((e11 = op1(e11, e11)) | ~ spl42_42), inference(avatar_component_clause, [], [f807])).
fof(f807, plain, (spl42_42 <=> (e11 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_42])])).
fof(f4512, plain, (~ spl42_52 | ~ spl42_60), inference(avatar_split_clause, [], [f4511, f883, f849])).
fof(f883, plain, (spl42_60 <=> (e13 = op1(e10, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_60])])).
fof(f4511, plain, (~ (e13 = op1(e10, e13)) | ~ spl42_60), inference(backward_demodulation, [], [f260, f885])).
fof(f885, plain, ((e13 = op1(e10, e11)) | ~ spl42_60), inference(avatar_component_clause, [], [f883])).
fof(f260, plain, ~ (op1(e10, e11) = op1(e10, e13)), inference(cnf_transformation, [], [f7])).
fof(f4507, plain, (~ spl42_69 | ~ spl42_70), inference(avatar_contradiction_clause, [], [f4506])).
fof(f4506, plain, ($false | (~ spl42_69 | ~ spl42_70)), inference(subsumption_resolution, [], [f4505, f334])).
fof(f334, plain, ~ (e20 = e21), inference(cnf_transformation, [], [f10])).
fof(f4505, plain, ((e20 = e21) | (~ spl42_69 | ~ spl42_70)), inference(backward_demodulation, [], [f960, f956])).
fof(f956, plain, ((e20 = op2(e23, e23)) | ~ spl42_69), inference(avatar_component_clause, [], [f954])).
fof(f954, plain, (spl42_69 <=> (e20 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_69])])).
fof(f960, plain, ((e21 = op2(e23, e23)) | ~ spl42_70), inference(avatar_component_clause, [], [f958])).
fof(f958, plain, (spl42_70 <=> (e21 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_70])])).
fof(f4494, plain, (~ spl42_89 | ~ spl42_91), inference(avatar_contradiction_clause, [], [f4493])).
fof(f4493, plain, ($false | (~ spl42_89 | ~ spl42_91)), inference(subsumption_resolution, [], [f4492, f335])).
fof(f335, plain, ~ (e20 = e22), inference(cnf_transformation, [], [f10])).
fof(f4492, plain, ((e20 = e22) | (~ spl42_89 | ~ spl42_91)), inference(backward_demodulation, [], [f1049, f1041])).
fof(f1049, plain, ((e22 = op2(e22, e22)) | ~ spl42_91), inference(avatar_component_clause, [], [f1047])).
fof(f1047, plain, (spl42_91 <=> (e22 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_91])])).
fof(f4453, plain, (spl42_91 | ~ spl42_135), inference(avatar_split_clause, [], [f4444, f1234, f1047])).
fof(f1234, plain, (spl42_135 <=> (e22 = unit2)), introduced(avatar_definition, [new_symbols(naming, [spl42_135])])).
fof(f4444, plain, ((e22 = op2(e22, e22)) | ~ spl42_135), inference(backward_demodulation, [], [f196, f1236])).
fof(f1236, plain, ((e22 = unit2) | ~ spl42_135), inference(avatar_component_clause, [], [f1234])).
fof(f196, plain, (e22 = op2(e22, unit2)), inference(cnf_transformation, [], [f5])).
fof(f4438, plain, (~ spl42_114 | ~ spl42_16 | ~ spl42_148 | spl42_232), inference(avatar_split_clause, [], [f4434, f2009, f1425, f696, f1145])).
fof(f2009, plain, (spl42_232 <=> (h11(op1(e13, e10)) = op2(e21, h10(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl42_232])])).
fof(f4434, plain, (~ (e21 = op2(e21, e20)) | (~ spl42_16 | ~ spl42_148 | spl42_232)), inference(backward_demodulation, [], [f4288, f1426])).
fof(f4288, plain, (~ (e21 = op2(e21, h10(e10))) | (~ spl42_16 | spl42_232)), inference(forward_demodulation, [], [f4287, f433])).
fof(f4287, plain, (~ (h11(e13) = op2(e21, h10(e10))) | (~ spl42_16 | spl42_232)), inference(forward_demodulation, [], [f2011, f698])).
fof(f698, plain, ((e13 = op1(e13, e10)) | ~ spl42_16), inference(avatar_component_clause, [], [f696])).
fof(f2011, plain, (~ (h11(op1(e13, e10)) = op2(e21, h10(e10))) | spl42_232), inference(avatar_component_clause, [], [f2009])).
fof(f4280, plain, (~ spl42_14 | ~ spl42_16), inference(avatar_contradiction_clause, [], [f4279])).
fof(f4279, plain, ($false | (~ spl42_14 | ~ spl42_16)), inference(subsumption_resolution, [], [f4278, f332])).
fof(f4278, plain, ((e11 = e13) | (~ spl42_14 | ~ spl42_16)), inference(forward_demodulation, [], [f698, f690])).
fof(f690, plain, ((e11 = op1(e13, e10)) | ~ spl42_14), inference(avatar_component_clause, [], [f688])).
fof(f688, plain, (spl42_14 <=> (e11 = op1(e13, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_14])])).
fof(f4260, plain, (~ spl42_51 | ~ spl42_55), inference(avatar_split_clause, [], [f4258, f862, f845])).
fof(f845, plain, (spl42_51 <=> (e12 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_51])])).
fof(f4258, plain, (~ (e12 = op1(e10, e13)) | ~ spl42_55), inference(backward_demodulation, [], [f261, f864])).
fof(f261, plain, ~ (op1(e10, e12) = op1(e10, e13)), inference(cnf_transformation, [], [f7])).
fof(f4249, plain, (~ spl42_21 | ~ spl42_67), inference(avatar_contradiction_clause, [], [f4248])).
fof(f4248, plain, ($false | (~ spl42_21 | ~ spl42_67)), inference(subsumption_resolution, [], [f4247, f329])).
fof(f329, plain, ~ (e10 = e12), inference(cnf_transformation, [], [f9])).
fof(f4247, plain, ((e10 = e12) | (~ spl42_21 | ~ spl42_67)), inference(forward_demodulation, [], [f4235, f720])).
fof(f4235, plain, ((e12 = op1(e12, e12)) | ~ spl42_67), inference(backward_demodulation, [], [f139, f915])).
fof(f915, plain, ((e12 = unit1) | ~ spl42_67), inference(avatar_component_clause, [], [f913])).
fof(f913, plain, (spl42_67 <=> (e12 = unit1)), introduced(avatar_definition, [new_symbols(naming, [spl42_67])])).
fof(f4209, plain, (~ spl42_109 | ~ spl42_110), inference(avatar_contradiction_clause, [], [f4208])).
fof(f4208, plain, ($false | (~ spl42_109 | ~ spl42_110)), inference(subsumption_resolution, [], [f4207, f334])).
fof(f4207, plain, ((e20 = e21) | (~ spl42_109 | ~ spl42_110)), inference(forward_demodulation, [], [f1130, f1126])).
fof(f1126, plain, ((e20 = op2(e21, e21)) | ~ spl42_109), inference(avatar_component_clause, [], [f1124])).
fof(f1124, plain, (spl42_109 <=> (e20 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_109])])).
fof(f1130, plain, ((e21 = op2(e21, e21)) | ~ spl42_110), inference(avatar_component_clause, [], [f1128])).
fof(f1128, plain, (spl42_110 <=> (e21 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_110])])).
fof(f4183, plain, (~ spl42_109 | ~ spl42_134), inference(avatar_contradiction_clause, [], [f4182])).
fof(f4182, plain, ($false | (~ spl42_109 | ~ spl42_134)), inference(subsumption_resolution, [], [f4181, f334])).
fof(f4181, plain, ((e20 = e21) | (~ spl42_109 | ~ spl42_134)), inference(forward_demodulation, [], [f4170, f1126])).
fof(f4170, plain, ((e21 = op2(e21, e21)) | ~ spl42_134), inference(backward_demodulation, [], [f194, f1232])).
fof(f1232, plain, ((e21 = unit2) | ~ spl42_134), inference(avatar_component_clause, [], [f1230])).
fof(f1230, plain, (spl42_134 <=> (e21 = unit2)), introduced(avatar_definition, [new_symbols(naming, [spl42_134])])).
fof(f4143, plain, (~ spl42_352 | ~ spl42_120), inference(avatar_split_clause, [], [f4142, f1170, f2954])).
fof(f4142, plain, (~ (e23 = h4(e12)) | ~ spl42_120), inference(forward_demodulation, [], [f1359, f1172])).
fof(f1359, plain, ~ (op2(e20, e23) = h4(e12)), inference(backward_demodulation, [], [f308, f407])).
fof(f308, plain, ~ (op2(e20, e21) = op2(e20, e23)), inference(cnf_transformation, [], [f8])).
fof(f4137, plain, (spl42_248 | ~ spl42_120), inference(avatar_split_clause, [], [f4136, f1170, f2151])).
fof(f4136, plain, ((e23 = h10(e12)) | ~ spl42_120), inference(forward_demodulation, [], [f431, f1172])).
fof(f4100, plain, (~ spl42_404 | ~ spl42_84), inference(avatar_split_clause, [], [f4099, f1017, f3346])).
fof(f4099, plain, (~ (e23 = h1(e12)) | ~ spl42_84), inference(forward_demodulation, [], [f1336, f1019])).
fof(f1336, plain, ~ (op2(e23, e20) = h1(e12)), inference(backward_demodulation, [], [f284, f395])).
fof(f284, plain, ~ (op2(e21, e20) = op2(e23, e20)), inference(cnf_transformation, [], [f8])).
fof(f4071, plain, (~ spl42_71 | ~ spl42_176), inference(avatar_contradiction_clause, [], [f4070])).
fof(f4070, plain, ($false | (~ spl42_71 | ~ spl42_176)), inference(subsumption_resolution, [], [f4069, f964])).
fof(f964, plain, ((e22 = op2(e23, e23)) | ~ spl42_71), inference(avatar_component_clause, [], [f962])).
fof(f962, plain, (spl42_71 <=> (e22 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_71])])).
fof(f4069, plain, (~ (e22 = op2(e23, e23)) | ~ spl42_176), inference(forward_demodulation, [], [f1370, f1598])).
fof(f1370, plain, ~ (op2(e23, e23) = h6(e12)), inference(backward_demodulation, [], [f326, f415])).
fof(f415, plain, (op2(e23, e21) = h6(e12)), inference(cnf_transformation, [], [f21])).
fof(f21, plain, ((op2(e23, e21) = h6(e12)) & (op2(e21, e21) = h6(e10)) & (e23 = h6(e13)) & (e21 = h6(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG037+1.p', ax21)).
fof(f326, plain, ~ (op2(e23, e21) = op2(e23, e23)), inference(cnf_transformation, [], [f8])).
fof(f3999, plain, (~ spl42_11 | ~ spl42_12), inference(avatar_contradiction_clause, [], [f3998])).
fof(f3998, plain, ($false | (~ spl42_11 | ~ spl42_12)), inference(subsumption_resolution, [], [f3997, f333])).
fof(f333, plain, ~ (e12 = e13), inference(cnf_transformation, [], [f9])).
fof(f3997, plain, ((e12 = e13) | (~ spl42_11 | ~ spl42_12)), inference(forward_demodulation, [], [f681, f677])).
fof(f681, plain, ((e13 = op1(e13, e11)) | ~ spl42_12), inference(avatar_component_clause, [], [f679])).
fof(f679, plain, (spl42_12 <=> (e13 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_12])])).
fof(f3953, plain, (~ spl42_41 | ~ spl42_66), inference(avatar_contradiction_clause, [], [f3952])).
fof(f3952, plain, ($false | (~ spl42_41 | ~ spl42_66)), inference(subsumption_resolution, [], [f3951, f328])).
fof(f3951, plain, ((e10 = e11) | (~ spl42_41 | ~ spl42_66)), inference(forward_demodulation, [], [f3941, f805])).
fof(f3941, plain, ((e11 = op1(e11, e11)) | ~ spl42_66), inference(backward_demodulation, [], [f137, f911])).
fof(f911, plain, ((e11 = unit1) | ~ spl42_66), inference(avatar_component_clause, [], [f909])).
fof(f909, plain, (spl42_66 <=> (e11 = unit1)), introduced(avatar_definition, [new_symbols(naming, [spl42_66])])).
fof(f3926, plain, (~ spl42_79 | ~ spl42_80), inference(avatar_contradiction_clause, [], [f3925])).
fof(f3925, plain, ($false | (~ spl42_79 | ~ spl42_80)), inference(subsumption_resolution, [], [f3924, f339])).
fof(f3924, plain, ((e22 = e23) | (~ spl42_79 | ~ spl42_80)), inference(forward_demodulation, [], [f1002, f998])).
fof(f998, plain, ((e22 = op2(e23, e21)) | ~ spl42_79), inference(avatar_component_clause, [], [f996])).
fof(f996, plain, (spl42_79 <=> (e22 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_79])])).
fof(f1002, plain, ((e23 = op2(e23, e21)) | ~ spl42_80), inference(avatar_component_clause, [], [f1000])).
fof(f1000, plain, (spl42_80 <=> (e23 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_80])])).
fof(f3922, plain, (~ spl42_82 | ~ spl42_84), inference(avatar_contradiction_clause, [], [f3921])).
fof(f3921, plain, ($false | (~ spl42_82 | ~ spl42_84)), inference(subsumption_resolution, [], [f3920, f338])).
fof(f338, plain, ~ (e21 = e23), inference(cnf_transformation, [], [f10])).
fof(f3920, plain, ((e21 = e23) | (~ spl42_82 | ~ spl42_84)), inference(forward_demodulation, [], [f1019, f1011])).
fof(f1011, plain, ((e21 = op2(e23, e20)) | ~ spl42_82), inference(avatar_component_clause, [], [f1009])).
fof(f1009, plain, (spl42_82 <=> (e21 = op2(e23, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_82])])).
fof(f3899, plain, (~ spl42_87 | ~ spl42_99), inference(avatar_contradiction_clause, [], [f3898])).
fof(f3898, plain, ($false | (~ spl42_87 | ~ spl42_99)), inference(subsumption_resolution, [], [f3895, f3520])).
fof(f3520, plain, (~ (e22 = h2(e12)) | ~ spl42_87), inference(backward_demodulation, [], [f1349, f1032])).
fof(f1032, plain, ((e22 = op2(e22, e23)) | ~ spl42_87), inference(avatar_component_clause, [], [f1030])).
fof(f1030, plain, (spl42_87 <=> (e22 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_87])])).
fof(f1349, plain, ~ (op2(e22, e23) = h2(e12)), inference(backward_demodulation, [], [f318, f399])).
fof(f399, plain, (op2(e22, e20) = h2(e12)), inference(cnf_transformation, [], [f17])).
fof(f17, plain, ((op2(e22, e20) = h2(e12)) & (op2(e20, e20) = h2(e10)) & (e22 = h2(e13)) & (e20 = h2(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG037+1.p', ax17)).
fof(f318, plain, ~ (op2(e22, e20) = op2(e22, e23)), inference(cnf_transformation, [], [f8])).
fof(f3895, plain, ((e22 = h2(e12)) | ~ spl42_99), inference(backward_demodulation, [], [f399, f1083])).
fof(f3881, plain, (~ spl42_109 | ~ spl42_111), inference(avatar_contradiction_clause, [], [f3880])).
fof(f3880, plain, ($false | (~ spl42_109 | ~ spl42_111)), inference(subsumption_resolution, [], [f3879, f335])).
fof(f3879, plain, ((e20 = e22) | (~ spl42_109 | ~ spl42_111)), inference(forward_demodulation, [], [f1134, f1126])).
fof(f1134, plain, ((e22 = op2(e21, e21)) | ~ spl42_111), inference(avatar_component_clause, [], [f1132])).
fof(f1132, plain, (spl42_111 <=> (e22 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_111])])).
fof(f3873, plain, (~ spl42_118 | ~ spl42_120), inference(avatar_contradiction_clause, [], [f3872])).
fof(f3872, plain, ($false | (~ spl42_118 | ~ spl42_120)), inference(subsumption_resolution, [], [f3871, f338])).
fof(f3871, plain, ((e21 = e23) | (~ spl42_118 | ~ spl42_120)), inference(forward_demodulation, [], [f1172, f1164])).
fof(f1164, plain, ((e21 = op2(e20, e23)) | ~ spl42_118), inference(avatar_component_clause, [], [f1162])).
fof(f1162, plain, (spl42_118 <=> (e21 = op2(e20, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_118])])).
fof(f3846, plain, (spl42_183 | ~ spl42_109), inference(avatar_split_clause, [], [f3845, f1124, f1639])).
fof(f3845, plain, ((e20 = h4(e10)) | ~ spl42_109), inference(forward_demodulation, [], [f406, f1126])).
fof(f406, plain, (op2(e21, e21) = h4(e10)), inference(cnf_transformation, [], [f19])).
fof(f3814, plain, (spl42_148 | ~ spl42_69), inference(avatar_split_clause, [], [f3813, f954, f1425])).
fof(f3813, plain, ((e20 = h10(e10)) | ~ spl42_69), inference(forward_demodulation, [], [f430, f956])).
fof(f3790, plain, (~ spl42_46 | ~ spl42_34), inference(avatar_split_clause, [], [f3789, f773, f824])).
fof(f773, plain, (spl42_34 <=> (e11 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_34])])).
fof(f3789, plain, (~ (e11 = op1(e11, e10)) | ~ spl42_34), inference(forward_demodulation, [], [f264, f775])).
fof(f775, plain, ((e11 = op1(e11, e13)) | ~ spl42_34), inference(avatar_component_clause, [], [f773])).
fof(f264, plain, ~ (op1(e11, e10) = op1(e11, e13)), inference(cnf_transformation, [], [f7])).
fof(f3784, plain, (~ spl42_37 | ~ spl42_41), inference(avatar_split_clause, [], [f3783, f803, f786])).
fof(f786, plain, (spl42_37 <=> (e10 = op1(e11, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_37])])).
fof(f3783, plain, (~ (e10 = op1(e11, e12)) | ~ spl42_41), inference(forward_demodulation, [], [f265, f805])).
fof(f265, plain, ~ (op1(e11, e11) = op1(e11, e12)), inference(cnf_transformation, [], [f7])).
fof(f3780, plain, (~ spl42_41 | spl42_238), inference(avatar_contradiction_clause, [], [f3779])).
fof(f3779, plain, ($false | (~ spl42_41 | spl42_238)), inference(subsumption_resolution, [], [f3778, f1392])).
fof(f3778, plain, (~ (h10(e10) = h11(e10)) | (~ spl42_41 | spl42_238)), inference(forward_demodulation, [], [f2035, f805])).
fof(f2035, plain, (~ (h10(e10) = h11(op1(e11, e11))) | spl42_238), inference(avatar_component_clause, [], [f2033])).
fof(f2033, plain, (spl42_238 <=> (h10(e10) = h11(op1(e11, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl42_238])])).
fof(f3730, plain, (~ spl42_10 | ~ spl42_11), inference(avatar_contradiction_clause, [], [f3729])).
fof(f3729, plain, ($false | (~ spl42_10 | ~ spl42_11)), inference(subsumption_resolution, [], [f3728, f331])).
fof(f331, plain, ~ (e11 = e12), inference(cnf_transformation, [], [f9])).
fof(f3728, plain, ((e11 = e12) | (~ spl42_10 | ~ spl42_11)), inference(backward_demodulation, [], [f677, f673])).
fof(f673, plain, ((e11 = op1(e13, e11)) | ~ spl42_10), inference(avatar_component_clause, [], [f671])).
fof(f3727, plain, (~ spl42_3 | ~ spl42_11), inference(avatar_split_clause, [], [f3725, f675, f641])).
fof(f641, plain, (spl42_3 <=> (e12 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_3])])).
fof(f3725, plain, (~ (e12 = op1(e13, e13)) | ~ spl42_11), inference(backward_demodulation, [], [f278, f677])).
fof(f278, plain, ~ (op1(e13, e11) = op1(e13, e13)), inference(cnf_transformation, [], [f7])).
fof(f3665, plain, (~ spl42_41 | ~ spl42_44), inference(avatar_contradiction_clause, [], [f3664])).
fof(f3664, plain, ($false | (~ spl42_41 | ~ spl42_44)), inference(subsumption_resolution, [], [f3663, f330])).
fof(f330, plain, ~ (e10 = e13), inference(cnf_transformation, [], [f9])).
fof(f3663, plain, ((e10 = e13) | (~ spl42_41 | ~ spl42_44)), inference(backward_demodulation, [], [f817, f805])).
fof(f817, plain, ((e13 = op1(e11, e11)) | ~ spl42_44), inference(avatar_component_clause, [], [f815])).
fof(f815, plain, (spl42_44 <=> (e13 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_44])])).
fof(f3593, plain, (spl42_52 | ~ spl42_65), inference(avatar_split_clause, [], [f3585, f905, f849])).
fof(f3585, plain, ((e13 = op1(e10, e13)) | ~ spl42_65), inference(backward_demodulation, [], [f140, f907])).
fof(f3591, plain, (spl42_55 | ~ spl42_65), inference(avatar_split_clause, [], [f3583, f905, f862])).
fof(f3583, plain, ((e12 = op1(e10, e12)) | ~ spl42_65), inference(backward_demodulation, [], [f138, f907])).
fof(f138, plain, (e12 = op1(unit1, e12)), inference(cnf_transformation, [], [f2])).
fof(f3554, plain, (~ spl42_78 | ~ spl42_79), inference(avatar_contradiction_clause, [], [f3553])).
fof(f3553, plain, ($false | (~ spl42_78 | ~ spl42_79)), inference(subsumption_resolution, [], [f3552, f337])).
fof(f337, plain, ~ (e21 = e22), inference(cnf_transformation, [], [f10])).
fof(f3552, plain, ((e21 = e22) | (~ spl42_78 | ~ spl42_79)), inference(backward_demodulation, [], [f998, f994])).
fof(f994, plain, ((e21 = op2(e23, e21)) | ~ spl42_78), inference(avatar_component_clause, [], [f992])).
fof(f3549, plain, (spl42_176 | ~ spl42_79), inference(avatar_split_clause, [], [f3545, f996, f1597])).
fof(f3545, plain, ((e22 = h6(e12)) | ~ spl42_79), inference(backward_demodulation, [], [f415, f998])).
fof(f3492, plain, (spl42_150 | ~ spl42_103), inference(avatar_split_clause, [], [f3486, f1098, f1435])).
fof(f1098, plain, (spl42_103 <=> (e22 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_103])])).
fof(f3486, plain, ((e22 = h11(e12)) | ~ spl42_103), inference(backward_demodulation, [], [f435, f1100])).
fof(f1100, plain, ((e22 = op2(e21, e23)) | ~ spl42_103), inference(avatar_component_clause, [], [f1098])).
fof(f435, plain, (op2(e21, e23) = h11(e12)), inference(cnf_transformation, [], [f26])).
fof(f3477, plain, (~ spl42_109 | ~ spl42_112), inference(avatar_contradiction_clause, [], [f3476])).
fof(f3476, plain, ($false | (~ spl42_109 | ~ spl42_112)), inference(subsumption_resolution, [], [f3475, f336])).
fof(f336, plain, ~ (e20 = e23), inference(cnf_transformation, [], [f10])).
fof(f3475, plain, ((e20 = e23) | (~ spl42_109 | ~ spl42_112)), inference(backward_demodulation, [], [f1138, f1126])).
fof(f1138, plain, ((e23 = op2(e21, e21)) | ~ spl42_112), inference(avatar_component_clause, [], [f1136])).
fof(f1136, plain, (spl42_112 <=> (e23 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_112])])).
fof(f3433, plain, (spl42_99 | ~ spl42_133), inference(avatar_split_clause, [], [f3425, f1226, f1081])).
fof(f3425, plain, ((e22 = op2(e22, e20)) | ~ spl42_133), inference(backward_demodulation, [], [f196, f1228])).
fof(f3432, plain, (spl42_123 | ~ spl42_133), inference(avatar_split_clause, [], [f3424, f1226, f1183])).
fof(f3424, plain, ((e22 = op2(e20, e22)) | ~ spl42_133), inference(backward_demodulation, [], [f195, f1228])).
fof(f195, plain, (e22 = op2(unit2, e22)), inference(cnf_transformation, [], [f5])).
fof(f2084, plain, (~ spl42_227 | spl42_152 | spl42_149 | ~ spl42_229 | ~ spl42_230 | ~ spl42_231 | ~ spl42_232 | ~ spl42_233 | ~ spl42_234 | ~ spl42_235 | ~ spl42_236 | ~ spl42_237 | ~ spl42_238 | ~ spl42_239 | ~ spl42_240 | ~ spl42_241 | ~ spl42_242 | ~ spl42_243), inference(avatar_split_clause, [], [f2083, f2053, f2049, f2045, f2041, f2037, f2033, f2029, f2025, f2021, f2017, f2013, f2009, f2005, f2001, f1997, f1431, f1449, f1989])).
fof(f1449, plain, (spl42_152 <=> sP36), introduced(avatar_definition, [new_symbols(naming, [spl42_152])])).
fof(f1431, plain, (spl42_149 <=> sP38), introduced(avatar_definition, [new_symbols(naming, [spl42_149])])).
fof(f2083, plain, (~ (op2(h10(e10), h10(e10)) = h11(op1(e10, e10))) | ~ (h11(op1(e10, e11)) = op2(h10(e10), e23)) | ~ (h11(op1(e10, e12)) = op2(h10(e10), h11(e12))) | ~ (h11(op1(e10, e13)) = op2(h10(e10), e21)) | ~ (h11(op1(e11, e10)) = op2(e23, h10(e10))) | ~ (h10(e10) = h11(op1(e11, e11))) | ~ (h11(op1(e11, e12)) = op2(e23, h11(e12))) | ~ (h6(e12) = h11(op1(e11, e13))) | ~ (h11(op1(e12, e10)) = op2(h11(e12), h10(e10))) | ~ (h11(op1(e12, e11)) = op2(h11(e12), e23)) | ~ (h11(op1(e12, e13)) = op2(h11(e12), e21)) | ~ (h11(op1(e13, e10)) = op2(e21, h10(e10))) | ~ (h11(e12) = h11(op1(e13, e11))) | ~ (h11(op1(e13, e12)) = op2(e21, h11(e12))) | ~ (h4(e10) = h11(op1(e13, e13))) | sP38 | sP36 | ~ (h11(op1(e12, e12)) = op2(h11(e12), h11(e12)))), inference(forward_demodulation, [], [f2082, f1392])).
fof(f2082, plain, (~ (h11(op1(e10, e11)) = op2(h10(e10), e23)) | ~ (h11(op1(e10, e12)) = op2(h10(e10), h11(e12))) | ~ (h11(op1(e10, e13)) = op2(h10(e10), e21)) | ~ (h11(op1(e11, e10)) = op2(e23, h10(e10))) | ~ (h10(e10) = h11(op1(e11, e11))) | ~ (h11(op1(e11, e12)) = op2(e23, h11(e12))) | ~ (h6(e12) = h11(op1(e11, e13))) | ~ (h11(op1(e12, e10)) = op2(h11(e12), h10(e10))) | ~ (h11(op1(e12, e11)) = op2(h11(e12), e23)) | ~ (h11(op1(e12, e13)) = op2(h11(e12), e21)) | ~ (h11(op1(e13, e10)) = op2(e21, h10(e10))) | ~ (h11(e12) = h11(op1(e13, e11))) | ~ (h11(op1(e13, e12)) = op2(e21, h11(e12))) | ~ (h4(e10) = h11(op1(e13, e13))) | sP38 | sP36 | ~ (h11(op1(e12, e12)) = op2(h11(e12), h11(e12))) | ~ (h11(op1(e10, e10)) = op2(h11(e10), h11(e10)))), inference(forward_demodulation, [], [f2081, f1392])).
fof(f2081, plain, (~ (h11(op1(e10, e11)) = op2(h11(e10), e23)) | ~ (h11(op1(e10, e12)) = op2(h10(e10), h11(e12))) | ~ (h11(op1(e10, e13)) = op2(h10(e10), e21)) | ~ (h11(op1(e11, e10)) = op2(e23, h10(e10))) | ~ (h10(e10) = h11(op1(e11, e11))) | ~ (h11(op1(e11, e12)) = op2(e23, h11(e12))) | ~ (h6(e12) = h11(op1(e11, e13))) | ~ (h11(op1(e12, e10)) = op2(h11(e12), h10(e10))) | ~ (h11(op1(e12, e11)) = op2(h11(e12), e23)) | ~ (h11(op1(e12, e13)) = op2(h11(e12), e21)) | ~ (h11(op1(e13, e10)) = op2(e21, h10(e10))) | ~ (h11(e12) = h11(op1(e13, e11))) | ~ (h11(op1(e13, e12)) = op2(e21, h11(e12))) | ~ (h4(e10) = h11(op1(e13, e13))) | sP38 | sP36 | ~ (h11(op1(e12, e12)) = op2(h11(e12), h11(e12))) | ~ (h11(op1(e10, e10)) = op2(h11(e10), h11(e10)))), inference(forward_demodulation, [], [f2080, f432])).
fof(f2080, plain, (~ (h11(op1(e10, e12)) = op2(h10(e10), h11(e12))) | ~ (h11(op1(e10, e13)) = op2(h10(e10), e21)) | ~ (h11(op1(e11, e10)) = op2(e23, h10(e10))) | ~ (h10(e10) = h11(op1(e11, e11))) | ~ (h11(op1(e11, e12)) = op2(e23, h11(e12))) | ~ (h6(e12) = h11(op1(e11, e13))) | ~ (h11(op1(e12, e10)) = op2(h11(e12), h10(e10))) | ~ (h11(op1(e12, e11)) = op2(h11(e12), e23)) | ~ (h11(op1(e12, e13)) = op2(h11(e12), e21)) | ~ (h11(op1(e13, e10)) = op2(e21, h10(e10))) | ~ (h11(e12) = h11(op1(e13, e11))) | ~ (h11(op1(e13, e12)) = op2(e21, h11(e12))) | ~ (h4(e10) = h11(op1(e13, e13))) | sP38 | sP36 | ~ (h11(op1(e12, e12)) = op2(h11(e12), h11(e12))) | ~ (h11(op1(e10, e11)) = op2(h11(e10), h11(e11))) | ~ (h11(op1(e10, e10)) = op2(h11(e10), h11(e10)))), inference(forward_demodulation, [], [f2079, f1392])).
fof(f2079, plain, (~ (h11(op1(e10, e13)) = op2(h10(e10), e21)) | ~ (h11(op1(e11, e10)) = op2(e23, h10(e10))) | ~ (h10(e10) = h11(op1(e11, e11))) | ~ (h11(op1(e11, e12)) = op2(e23, h11(e12))) | ~ (h6(e12) = h11(op1(e11, e13))) | ~ (h11(op1(e12, e10)) = op2(h11(e12), h10(e10))) | ~ (h11(op1(e12, e11)) = op2(h11(e12), e23)) | ~ (h11(op1(e12, e13)) = op2(h11(e12), e21)) | ~ (h11(op1(e13, e10)) = op2(e21, h10(e10))) | ~ (h11(e12) = h11(op1(e13, e11))) | ~ (h11(op1(e13, e12)) = op2(e21, h11(e12))) | ~ (h4(e10) = h11(op1(e13, e13))) | sP38 | sP36 | ~ (h11(op1(e12, e12)) = op2(h11(e12), h11(e12))) | ~ (h11(op1(e10, e12)) = op2(h11(e10), h11(e12))) | ~ (h11(op1(e10, e11)) = op2(h11(e10), h11(e11))) | ~ (h11(op1(e10, e10)) = op2(h11(e10), h11(e10)))), inference(forward_demodulation, [], [f2078, f1392])).
fof(f2078, plain, (~ (h11(op1(e10, e13)) = op2(h11(e10), e21)) | ~ (h11(op1(e11, e10)) = op2(e23, h10(e10))) | ~ (h10(e10) = h11(op1(e11, e11))) | ~ (h11(op1(e11, e12)) = op2(e23, h11(e12))) | ~ (h6(e12) = h11(op1(e11, e13))) | ~ (h11(op1(e12, e10)) = op2(h11(e12), h10(e10))) | ~ (h11(op1(e12, e11)) = op2(h11(e12), e23)) | ~ (h11(op1(e12, e13)) = op2(h11(e12), e21)) | ~ (h11(op1(e13, e10)) = op2(e21, h10(e10))) | ~ (h11(e12) = h11(op1(e13, e11))) | ~ (h11(op1(e13, e12)) = op2(e21, h11(e12))) | ~ (h4(e10) = h11(op1(e13, e13))) | sP38 | sP36 | ~ (h11(op1(e12, e12)) = op2(h11(e12), h11(e12))) | ~ (h11(op1(e10, e12)) = op2(h11(e10), h11(e12))) | ~ (h11(op1(e10, e11)) = op2(h11(e10), h11(e11))) | ~ (h11(op1(e10, e10)) = op2(h11(e10), h11(e10)))), inference(forward_demodulation, [], [f2077, f433])).
fof(f2077, plain, (~ (h11(op1(e11, e10)) = op2(e23, h10(e10))) | ~ (h10(e10) = h11(op1(e11, e11))) | ~ (h11(op1(e11, e12)) = op2(e23, h11(e12))) | ~ (h6(e12) = h11(op1(e11, e13))) | ~ (h11(op1(e12, e10)) = op2(h11(e12), h10(e10))) | ~ (h11(op1(e12, e11)) = op2(h11(e12), e23)) | ~ (h11(op1(e12, e13)) = op2(h11(e12), e21)) | ~ (h11(op1(e13, e10)) = op2(e21, h10(e10))) | ~ (h11(e12) = h11(op1(e13, e11))) | ~ (h11(op1(e13, e12)) = op2(e21, h11(e12))) | ~ (h4(e10) = h11(op1(e13, e13))) | sP38 | sP36 | ~ (h11(op1(e12, e12)) = op2(h11(e12), h11(e12))) | ~ (h11(op1(e10, e13)) = op2(h11(e10), h11(e13))) | ~ (h11(op1(e10, e12)) = op2(h11(e10), h11(e12))) | ~ (h11(op1(e10, e11)) = op2(h11(e10), h11(e11))) | ~ (h11(op1(e10, e10)) = op2(h11(e10), h11(e10)))), inference(forward_demodulation, [], [f2076, f432])).
fof(f2076, plain, (~ (h11(op1(e11, e10)) = op2(h11(e11), h10(e10))) | ~ (h10(e10) = h11(op1(e11, e11))) | ~ (h11(op1(e11, e12)) = op2(e23, h11(e12))) | ~ (h6(e12) = h11(op1(e11, e13))) | ~ (h11(op1(e12, e10)) = op2(h11(e12), h10(e10))) | ~ (h11(op1(e12, e11)) = op2(h11(e12), e23)) | ~ (h11(op1(e12, e13)) = op2(h11(e12), e21)) | ~ (h11(op1(e13, e10)) = op2(e21, h10(e10))) | ~ (h11(e12) = h11(op1(e13, e11))) | ~ (h11(op1(e13, e12)) = op2(e21, h11(e12))) | ~ (h4(e10) = h11(op1(e13, e13))) | sP38 | sP36 | ~ (h11(op1(e12, e12)) = op2(h11(e12), h11(e12))) | ~ (h11(op1(e10, e13)) = op2(h11(e10), h11(e13))) | ~ (h11(op1(e10, e12)) = op2(h11(e10), h11(e12))) | ~ (h11(op1(e10, e11)) = op2(h11(e10), h11(e11))) | ~ (h11(op1(e10, e10)) = op2(h11(e10), h11(e10)))), inference(forward_demodulation, [], [f2075, f1392])).
fof(f2075, plain, (~ (h10(e10) = h11(op1(e11, e11))) | ~ (h11(op1(e11, e12)) = op2(e23, h11(e12))) | ~ (h6(e12) = h11(op1(e11, e13))) | ~ (h11(op1(e12, e10)) = op2(h11(e12), h10(e10))) | ~ (h11(op1(e12, e11)) = op2(h11(e12), e23)) | ~ (h11(op1(e12, e13)) = op2(h11(e12), e21)) | ~ (h11(op1(e13, e10)) = op2(e21, h10(e10))) | ~ (h11(e12) = h11(op1(e13, e11))) | ~ (h11(op1(e13, e12)) = op2(e21, h11(e12))) | ~ (h4(e10) = h11(op1(e13, e13))) | sP38 | sP36 | ~ (h11(op1(e12, e12)) = op2(h11(e12), h11(e12))) | ~ (h11(op1(e11, e10)) = op2(h11(e11), h11(e10))) | ~ (h11(op1(e10, e13)) = op2(h11(e10), h11(e13))) | ~ (h11(op1(e10, e12)) = op2(h11(e10), h11(e12))) | ~ (h11(op1(e10, e11)) = op2(h11(e10), h11(e11))) | ~ (h11(op1(e10, e10)) = op2(h11(e10), h11(e10)))), inference(forward_demodulation, [], [f2074, f430])).
fof(f2074, plain, (~ (op2(e23, e23) = h11(op1(e11, e11))) | ~ (h11(op1(e11, e12)) = op2(e23, h11(e12))) | ~ (h6(e12) = h11(op1(e11, e13))) | ~ (h11(op1(e12, e10)) = op2(h11(e12), h10(e10))) | ~ (h11(op1(e12, e11)) = op2(h11(e12), e23)) | ~ (h11(op1(e12, e13)) = op2(h11(e12), e21)) | ~ (h11(op1(e13, e10)) = op2(e21, h10(e10))) | ~ (h11(e12) = h11(op1(e13, e11))) | ~ (h11(op1(e13, e12)) = op2(e21, h11(e12))) | ~ (h4(e10) = h11(op1(e13, e13))) | sP38 | sP36 | ~ (h11(op1(e12, e12)) = op2(h11(e12), h11(e12))) | ~ (h11(op1(e11, e10)) = op2(h11(e11), h11(e10))) | ~ (h11(op1(e10, e13)) = op2(h11(e10), h11(e13))) | ~ (h11(op1(e10, e12)) = op2(h11(e10), h11(e12))) | ~ (h11(op1(e10, e11)) = op2(h11(e10), h11(e11))) | ~ (h11(op1(e10, e10)) = op2(h11(e10), h11(e10)))), inference(forward_demodulation, [], [f2073, f432])).
fof(f2073, plain, (~ (h11(op1(e11, e12)) = op2(e23, h11(e12))) | ~ (h6(e12) = h11(op1(e11, e13))) | ~ (h11(op1(e12, e10)) = op2(h11(e12), h10(e10))) | ~ (h11(op1(e12, e11)) = op2(h11(e12), e23)) | ~ (h11(op1(e12, e13)) = op2(h11(e12), e21)) | ~ (h11(op1(e13, e10)) = op2(e21, h10(e10))) | ~ (h11(e12) = h11(op1(e13, e11))) | ~ (h11(op1(e13, e12)) = op2(e21, h11(e12))) | ~ (h4(e10) = h11(op1(e13, e13))) | sP38 | sP36 | ~ (h11(op1(e12, e12)) = op2(h11(e12), h11(e12))) | ~ (h11(op1(e11, e11)) = op2(h11(e11), h11(e11))) | ~ (h11(op1(e11, e10)) = op2(h11(e11), h11(e10))) | ~ (h11(op1(e10, e13)) = op2(h11(e10), h11(e13))) | ~ (h11(op1(e10, e12)) = op2(h11(e10), h11(e12))) | ~ (h11(op1(e10, e11)) = op2(h11(e10), h11(e11))) | ~ (h11(op1(e10, e10)) = op2(h11(e10), h11(e10)))), inference(forward_demodulation, [], [f2072, f432])).
fof(f2072, plain, (~ (h6(e12) = h11(op1(e11, e13))) | ~ (h11(op1(e12, e10)) = op2(h11(e12), h10(e10))) | ~ (h11(op1(e12, e11)) = op2(h11(e12), e23)) | ~ (h11(op1(e12, e13)) = op2(h11(e12), e21)) | ~ (h11(op1(e13, e10)) = op2(e21, h10(e10))) | ~ (h11(e12) = h11(op1(e13, e11))) | ~ (h11(op1(e13, e12)) = op2(e21, h11(e12))) | ~ (h4(e10) = h11(op1(e13, e13))) | sP38 | sP36 | ~ (h11(op1(e12, e12)) = op2(h11(e12), h11(e12))) | ~ (h11(op1(e11, e12)) = op2(h11(e11), h11(e12))) | ~ (h11(op1(e11, e11)) = op2(h11(e11), h11(e11))) | ~ (h11(op1(e11, e10)) = op2(h11(e11), h11(e10))) | ~ (h11(op1(e10, e13)) = op2(h11(e10), h11(e13))) | ~ (h11(op1(e10, e12)) = op2(h11(e10), h11(e12))) | ~ (h11(op1(e10, e11)) = op2(h11(e10), h11(e11))) | ~ (h11(op1(e10, e10)) = op2(h11(e10), h11(e10)))), inference(forward_demodulation, [], [f2071, f415])).
fof(f2071, plain, (~ (op2(e23, e21) = h11(op1(e11, e13))) | ~ (h11(op1(e12, e10)) = op2(h11(e12), h10(e10))) | ~ (h11(op1(e12, e11)) = op2(h11(e12), e23)) | ~ (h11(op1(e12, e13)) = op2(h11(e12), e21)) | ~ (h11(op1(e13, e10)) = op2(e21, h10(e10))) | ~ (h11(e12) = h11(op1(e13, e11))) | ~ (h11(op1(e13, e12)) = op2(e21, h11(e12))) | ~ (h4(e10) = h11(op1(e13, e13))) | sP38 | sP36 | ~ (h11(op1(e12, e12)) = op2(h11(e12), h11(e12))) | ~ (h11(op1(e11, e12)) = op2(h11(e11), h11(e12))) | ~ (h11(op1(e11, e11)) = op2(h11(e11), h11(e11))) | ~ (h11(op1(e11, e10)) = op2(h11(e11), h11(e10))) | ~ (h11(op1(e10, e13)) = op2(h11(e10), h11(e13))) | ~ (h11(op1(e10, e12)) = op2(h11(e10), h11(e12))) | ~ (h11(op1(e10, e11)) = op2(h11(e10), h11(e11))) | ~ (h11(op1(e10, e10)) = op2(h11(e10), h11(e10)))), inference(forward_demodulation, [], [f2070, f432])).
fof(f2070, plain, (~ (h11(op1(e11, e13)) = op2(h11(e11), e21)) | ~ (h11(op1(e12, e10)) = op2(h11(e12), h10(e10))) | ~ (h11(op1(e12, e11)) = op2(h11(e12), e23)) | ~ (h11(op1(e12, e13)) = op2(h11(e12), e21)) | ~ (h11(op1(e13, e10)) = op2(e21, h10(e10))) | ~ (h11(e12) = h11(op1(e13, e11))) | ~ (h11(op1(e13, e12)) = op2(e21, h11(e12))) | ~ (h4(e10) = h11(op1(e13, e13))) | sP38 | sP36 | ~ (h11(op1(e12, e12)) = op2(h11(e12), h11(e12))) | ~ (h11(op1(e11, e12)) = op2(h11(e11), h11(e12))) | ~ (h11(op1(e11, e11)) = op2(h11(e11), h11(e11))) | ~ (h11(op1(e11, e10)) = op2(h11(e11), h11(e10))) | ~ (h11(op1(e10, e13)) = op2(h11(e10), h11(e13))) | ~ (h11(op1(e10, e12)) = op2(h11(e10), h11(e12))) | ~ (h11(op1(e10, e11)) = op2(h11(e10), h11(e11))) | ~ (h11(op1(e10, e10)) = op2(h11(e10), h11(e10)))), inference(forward_demodulation, [], [f2069, f433])).
fof(f2069, plain, (~ (h11(op1(e12, e10)) = op2(h11(e12), h10(e10))) | ~ (h11(op1(e12, e11)) = op2(h11(e12), e23)) | ~ (h11(op1(e12, e13)) = op2(h11(e12), e21)) | ~ (h11(op1(e13, e10)) = op2(e21, h10(e10))) | ~ (h11(e12) = h11(op1(e13, e11))) | ~ (h11(op1(e13, e12)) = op2(e21, h11(e12))) | ~ (h4(e10) = h11(op1(e13, e13))) | sP38 | sP36 | ~ (h11(op1(e12, e12)) = op2(h11(e12), h11(e12))) | ~ (h11(op1(e11, e13)) = op2(h11(e11), h11(e13))) | ~ (h11(op1(e11, e12)) = op2(h11(e11), h11(e12))) | ~ (h11(op1(e11, e11)) = op2(h11(e11), h11(e11))) | ~ (h11(op1(e11, e10)) = op2(h11(e11), h11(e10))) | ~ (h11(op1(e10, e13)) = op2(h11(e10), h11(e13))) | ~ (h11(op1(e10, e12)) = op2(h11(e10), h11(e12))) | ~ (h11(op1(e10, e11)) = op2(h11(e10), h11(e11))) | ~ (h11(op1(e10, e10)) = op2(h11(e10), h11(e10)))), inference(forward_demodulation, [], [f2068, f1392])).
fof(f2068, plain, (~ (h11(op1(e12, e11)) = op2(h11(e12), e23)) | ~ (h11(op1(e12, e13)) = op2(h11(e12), e21)) | ~ (h11(op1(e13, e10)) = op2(e21, h10(e10))) | ~ (h11(e12) = h11(op1(e13, e11))) | ~ (h11(op1(e13, e12)) = op2(e21, h11(e12))) | ~ (h4(e10) = h11(op1(e13, e13))) | sP38 | sP36 | ~ (h11(op1(e12, e12)) = op2(h11(e12), h11(e12))) | ~ (h11(op1(e12, e10)) = op2(h11(e12), h11(e10))) | ~ (h11(op1(e11, e13)) = op2(h11(e11), h11(e13))) | ~ (h11(op1(e11, e12)) = op2(h11(e11), h11(e12))) | ~ (h11(op1(e11, e11)) = op2(h11(e11), h11(e11))) | ~ (h11(op1(e11, e10)) = op2(h11(e11), h11(e10))) | ~ (h11(op1(e10, e13)) = op2(h11(e10), h11(e13))) | ~ (h11(op1(e10, e12)) = op2(h11(e10), h11(e12))) | ~ (h11(op1(e10, e11)) = op2(h11(e10), h11(e11))) | ~ (h11(op1(e10, e10)) = op2(h11(e10), h11(e10)))), inference(forward_demodulation, [], [f2067, f432])).
fof(f2067, plain, (~ (h11(op1(e12, e13)) = op2(h11(e12), e21)) | ~ (h11(op1(e13, e10)) = op2(e21, h10(e10))) | ~ (h11(e12) = h11(op1(e13, e11))) | ~ (h11(op1(e13, e12)) = op2(e21, h11(e12))) | ~ (h4(e10) = h11(op1(e13, e13))) | sP38 | sP36 | ~ (h11(op1(e12, e12)) = op2(h11(e12), h11(e12))) | ~ (h11(op1(e12, e11)) = op2(h11(e12), h11(e11))) | ~ (h11(op1(e12, e10)) = op2(h11(e12), h11(e10))) | ~ (h11(op1(e11, e13)) = op2(h11(e11), h11(e13))) | ~ (h11(op1(e11, e12)) = op2(h11(e11), h11(e12))) | ~ (h11(op1(e11, e11)) = op2(h11(e11), h11(e11))) | ~ (h11(op1(e11, e10)) = op2(h11(e11), h11(e10))) | ~ (h11(op1(e10, e13)) = op2(h11(e10), h11(e13))) | ~ (h11(op1(e10, e12)) = op2(h11(e10), h11(e12))) | ~ (h11(op1(e10, e11)) = op2(h11(e10), h11(e11))) | ~ (h11(op1(e10, e10)) = op2(h11(e10), h11(e10)))), inference(forward_demodulation, [], [f2066, f433])).
fof(f2066, plain, (~ (h11(op1(e13, e10)) = op2(e21, h10(e10))) | ~ (h11(e12) = h11(op1(e13, e11))) | ~ (h11(op1(e13, e12)) = op2(e21, h11(e12))) | ~ (h4(e10) = h11(op1(e13, e13))) | sP38 | sP36 | ~ (h11(op1(e12, e13)) = op2(h11(e12), h11(e13))) | ~ (h11(op1(e12, e12)) = op2(h11(e12), h11(e12))) | ~ (h11(op1(e12, e11)) = op2(h11(e12), h11(e11))) | ~ (h11(op1(e12, e10)) = op2(h11(e12), h11(e10))) | ~ (h11(op1(e11, e13)) = op2(h11(e11), h11(e13))) | ~ (h11(op1(e11, e12)) = op2(h11(e11), h11(e12))) | ~ (h11(op1(e11, e11)) = op2(h11(e11), h11(e11))) | ~ (h11(op1(e11, e10)) = op2(h11(e11), h11(e10))) | ~ (h11(op1(e10, e13)) = op2(h11(e10), h11(e13))) | ~ (h11(op1(e10, e12)) = op2(h11(e10), h11(e12))) | ~ (h11(op1(e10, e11)) = op2(h11(e10), h11(e11))) | ~ (h11(op1(e10, e10)) = op2(h11(e10), h11(e10)))), inference(forward_demodulation, [], [f2065, f433])).
fof(f2065, plain, (~ (h11(op1(e13, e10)) = op2(h11(e13), h10(e10))) | ~ (h11(e12) = h11(op1(e13, e11))) | ~ (h11(op1(e13, e12)) = op2(e21, h11(e12))) | ~ (h4(e10) = h11(op1(e13, e13))) | sP38 | sP36 | ~ (h11(op1(e12, e13)) = op2(h11(e12), h11(e13))) | ~ (h11(op1(e12, e12)) = op2(h11(e12), h11(e12))) | ~ (h11(op1(e12, e11)) = op2(h11(e12), h11(e11))) | ~ (h11(op1(e12, e10)) = op2(h11(e12), h11(e10))) | ~ (h11(op1(e11, e13)) = op2(h11(e11), h11(e13))) | ~ (h11(op1(e11, e12)) = op2(h11(e11), h11(e12))) | ~ (h11(op1(e11, e11)) = op2(h11(e11), h11(e11))) | ~ (h11(op1(e11, e10)) = op2(h11(e11), h11(e10))) | ~ (h11(op1(e10, e13)) = op2(h11(e10), h11(e13))) | ~ (h11(op1(e10, e12)) = op2(h11(e10), h11(e12))) | ~ (h11(op1(e10, e11)) = op2(h11(e10), h11(e11))) | ~ (h11(op1(e10, e10)) = op2(h11(e10), h11(e10)))), inference(forward_demodulation, [], [f2064, f1392])).
fof(f2064, plain, (~ (h11(e12) = h11(op1(e13, e11))) | ~ (h11(op1(e13, e12)) = op2(e21, h11(e12))) | ~ (h4(e10) = h11(op1(e13, e13))) | sP38 | sP36 | ~ (h11(op1(e13, e10)) = op2(h11(e13), h11(e10))) | ~ (h11(op1(e12, e13)) = op2(h11(e12), h11(e13))) | ~ (h11(op1(e12, e12)) = op2(h11(e12), h11(e12))) | ~ (h11(op1(e12, e11)) = op2(h11(e12), h11(e11))) | ~ (h11(op1(e12, e10)) = op2(h11(e12), h11(e10))) | ~ (h11(op1(e11, e13)) = op2(h11(e11), h11(e13))) | ~ (h11(op1(e11, e12)) = op2(h11(e11), h11(e12))) | ~ (h11(op1(e11, e11)) = op2(h11(e11), h11(e11))) | ~ (h11(op1(e11, e10)) = op2(h11(e11), h11(e10))) | ~ (h11(op1(e10, e13)) = op2(h11(e10), h11(e13))) | ~ (h11(op1(e10, e12)) = op2(h11(e10), h11(e12))) | ~ (h11(op1(e10, e11)) = op2(h11(e10), h11(e11))) | ~ (h11(op1(e10, e10)) = op2(h11(e10), h11(e10)))), inference(forward_demodulation, [], [f2063, f435])).
fof(f2063, plain, (~ (op2(e21, e23) = h11(op1(e13, e11))) | ~ (h11(op1(e13, e12)) = op2(e21, h11(e12))) | ~ (h4(e10) = h11(op1(e13, e13))) | sP38 | sP36 | ~ (h11(op1(e13, e10)) = op2(h11(e13), h11(e10))) | ~ (h11(op1(e12, e13)) = op2(h11(e12), h11(e13))) | ~ (h11(op1(e12, e12)) = op2(h11(e12), h11(e12))) | ~ (h11(op1(e12, e11)) = op2(h11(e12), h11(e11))) | ~ (h11(op1(e12, e10)) = op2(h11(e12), h11(e10))) | ~ (h11(op1(e11, e13)) = op2(h11(e11), h11(e13))) | ~ (h11(op1(e11, e12)) = op2(h11(e11), h11(e12))) | ~ (h11(op1(e11, e11)) = op2(h11(e11), h11(e11))) | ~ (h11(op1(e11, e10)) = op2(h11(e11), h11(e10))) | ~ (h11(op1(e10, e13)) = op2(h11(e10), h11(e13))) | ~ (h11(op1(e10, e12)) = op2(h11(e10), h11(e12))) | ~ (h11(op1(e10, e11)) = op2(h11(e10), h11(e11))) | ~ (h11(op1(e10, e10)) = op2(h11(e10), h11(e10)))), inference(forward_demodulation, [], [f2062, f433])).
fof(f2062, plain, (~ (h11(op1(e13, e11)) = op2(h11(e13), e23)) | ~ (h11(op1(e13, e12)) = op2(e21, h11(e12))) | ~ (h4(e10) = h11(op1(e13, e13))) | sP38 | sP36 | ~ (h11(op1(e13, e10)) = op2(h11(e13), h11(e10))) | ~ (h11(op1(e12, e13)) = op2(h11(e12), h11(e13))) | ~ (h11(op1(e12, e12)) = op2(h11(e12), h11(e12))) | ~ (h11(op1(e12, e11)) = op2(h11(e12), h11(e11))) | ~ (h11(op1(e12, e10)) = op2(h11(e12), h11(e10))) | ~ (h11(op1(e11, e13)) = op2(h11(e11), h11(e13))) | ~ (h11(op1(e11, e12)) = op2(h11(e11), h11(e12))) | ~ (h11(op1(e11, e11)) = op2(h11(e11), h11(e11))) | ~ (h11(op1(e11, e10)) = op2(h11(e11), h11(e10))) | ~ (h11(op1(e10, e13)) = op2(h11(e10), h11(e13))) | ~ (h11(op1(e10, e12)) = op2(h11(e10), h11(e12))) | ~ (h11(op1(e10, e11)) = op2(h11(e10), h11(e11))) | ~ (h11(op1(e10, e10)) = op2(h11(e10), h11(e10)))), inference(forward_demodulation, [], [f2061, f432])).
fof(f2061, plain, (~ (h11(op1(e13, e12)) = op2(e21, h11(e12))) | ~ (h4(e10) = h11(op1(e13, e13))) | sP38 | sP36 | ~ (h11(op1(e13, e11)) = op2(h11(e13), h11(e11))) | ~ (h11(op1(e13, e10)) = op2(h11(e13), h11(e10))) | ~ (h11(op1(e12, e13)) = op2(h11(e12), h11(e13))) | ~ (h11(op1(e12, e12)) = op2(h11(e12), h11(e12))) | ~ (h11(op1(e12, e11)) = op2(h11(e12), h11(e11))) | ~ (h11(op1(e12, e10)) = op2(h11(e12), h11(e10))) | ~ (h11(op1(e11, e13)) = op2(h11(e11), h11(e13))) | ~ (h11(op1(e11, e12)) = op2(h11(e11), h11(e12))) | ~ (h11(op1(e11, e11)) = op2(h11(e11), h11(e11))) | ~ (h11(op1(e11, e10)) = op2(h11(e11), h11(e10))) | ~ (h11(op1(e10, e13)) = op2(h11(e10), h11(e13))) | ~ (h11(op1(e10, e12)) = op2(h11(e10), h11(e12))) | ~ (h11(op1(e10, e11)) = op2(h11(e10), h11(e11))) | ~ (h11(op1(e10, e10)) = op2(h11(e10), h11(e10)))), inference(forward_demodulation, [], [f2060, f433])).
fof(f2060, plain, (~ (h4(e10) = h11(op1(e13, e13))) | sP38 | sP36 | ~ (h11(op1(e13, e12)) = op2(h11(e13), h11(e12))) | ~ (h11(op1(e13, e11)) = op2(h11(e13), h11(e11))) | ~ (h11(op1(e13, e10)) = op2(h11(e13), h11(e10))) | ~ (h11(op1(e12, e13)) = op2(h11(e12), h11(e13))) | ~ (h11(op1(e12, e12)) = op2(h11(e12), h11(e12))) | ~ (h11(op1(e12, e11)) = op2(h11(e12), h11(e11))) | ~ (h11(op1(e12, e10)) = op2(h11(e12), h11(e10))) | ~ (h11(op1(e11, e13)) = op2(h11(e11), h11(e13))) | ~ (h11(op1(e11, e12)) = op2(h11(e11), h11(e12))) | ~ (h11(op1(e11, e11)) = op2(h11(e11), h11(e11))) | ~ (h11(op1(e11, e10)) = op2(h11(e11), h11(e10))) | ~ (h11(op1(e10, e13)) = op2(h11(e10), h11(e13))) | ~ (h11(op1(e10, e12)) = op2(h11(e10), h11(e12))) | ~ (h11(op1(e10, e11)) = op2(h11(e10), h11(e11))) | ~ (h11(op1(e10, e10)) = op2(h11(e10), h11(e10)))), inference(forward_demodulation, [], [f2059, f406])).
fof(f2059, plain, (~ (op2(e21, e21) = h11(op1(e13, e13))) | sP38 | sP36 | ~ (h11(op1(e13, e12)) = op2(h11(e13), h11(e12))) | ~ (h11(op1(e13, e11)) = op2(h11(e13), h11(e11))) | ~ (h11(op1(e13, e10)) = op2(h11(e13), h11(e10))) | ~ (h11(op1(e12, e13)) = op2(h11(e12), h11(e13))) | ~ (h11(op1(e12, e12)) = op2(h11(e12), h11(e12))) | ~ (h11(op1(e12, e11)) = op2(h11(e12), h11(e11))) | ~ (h11(op1(e12, e10)) = op2(h11(e12), h11(e10))) | ~ (h11(op1(e11, e13)) = op2(h11(e11), h11(e13))) | ~ (h11(op1(e11, e12)) = op2(h11(e11), h11(e12))) | ~ (h11(op1(e11, e11)) = op2(h11(e11), h11(e11))) | ~ (h11(op1(e11, e10)) = op2(h11(e11), h11(e10))) | ~ (h11(op1(e10, e13)) = op2(h11(e10), h11(e13))) | ~ (h11(op1(e10, e12)) = op2(h11(e10), h11(e12))) | ~ (h11(op1(e10, e11)) = op2(h11(e10), h11(e11))) | ~ (h11(op1(e10, e10)) = op2(h11(e10), h11(e10)))), inference(forward_demodulation, [], [f2058, f433])).
fof(f2058, plain, (sP38 | sP36 | ~ (h11(op1(e13, e13)) = op2(h11(e13), h11(e13))) | ~ (h11(op1(e13, e12)) = op2(h11(e13), h11(e12))) | ~ (h11(op1(e13, e11)) = op2(h11(e13), h11(e11))) | ~ (h11(op1(e13, e10)) = op2(h11(e13), h11(e10))) | ~ (h11(op1(e12, e13)) = op2(h11(e12), h11(e13))) | ~ (h11(op1(e12, e12)) = op2(h11(e12), h11(e12))) | ~ (h11(op1(e12, e11)) = op2(h11(e12), h11(e11))) | ~ (h11(op1(e12, e10)) = op2(h11(e12), h11(e10))) | ~ (h11(op1(e11, e13)) = op2(h11(e11), h11(e13))) | ~ (h11(op1(e11, e12)) = op2(h11(e11), h11(e12))) | ~ (h11(op1(e11, e11)) = op2(h11(e11), h11(e11))) | ~ (h11(op1(e11, e10)) = op2(h11(e11), h11(e10))) | ~ (h11(op1(e10, e13)) = op2(h11(e10), h11(e13))) | ~ (h11(op1(e10, e12)) = op2(h11(e10), h11(e12))) | ~ (h11(op1(e10, e11)) = op2(h11(e10), h11(e11))) | ~ (h11(op1(e10, e10)) = op2(h11(e10), h11(e10)))), inference(subsumption_resolution, [], [f2057, f1446])).
fof(f1446, plain, ~ sP37, inference(subsumption_resolution, [], [f459, f433])).
fof(f459, plain, (~ (e21 = h11(e13)) | ~ sP37), inference(cnf_transformation, [], [f86])).
fof(f86, plain, ((~ (e21 = h11(e13)) & ~ (e21 = h11(e12)) & ~ (e21 = h11(e11)) & ~ (e21 = h11(e10))) | ~ sP37), inference(nnf_transformation, [], [f70])).
fof(f70, plain, ((~ (e21 = h11(e13)) & ~ (e21 = h11(e12)) & ~ (e21 = h11(e11)) & ~ (e21 = h11(e10))) | ~ sP37), inference(usedef, [], [e70])).
fof(e70, plain, (sP37 <=> (~ (e21 = h11(e13)) & ~ (e21 = h11(e12)) & ~ (e21 = h11(e11)) & ~ (e21 = h11(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP37])])).
fof(f2057, plain, (sP38 | sP37 | sP36 | ~ (h11(op1(e13, e13)) = op2(h11(e13), h11(e13))) | ~ (h11(op1(e13, e12)) = op2(h11(e13), h11(e12))) | ~ (h11(op1(e13, e11)) = op2(h11(e13), h11(e11))) | ~ (h11(op1(e13, e10)) = op2(h11(e13), h11(e10))) | ~ (h11(op1(e12, e13)) = op2(h11(e12), h11(e13))) | ~ (h11(op1(e12, e12)) = op2(h11(e12), h11(e12))) | ~ (h11(op1(e12, e11)) = op2(h11(e12), h11(e11))) | ~ (h11(op1(e12, e10)) = op2(h11(e12), h11(e10))) | ~ (h11(op1(e11, e13)) = op2(h11(e11), h11(e13))) | ~ (h11(op1(e11, e12)) = op2(h11(e11), h11(e12))) | ~ (h11(op1(e11, e11)) = op2(h11(e11), h11(e11))) | ~ (h11(op1(e11, e10)) = op2(h11(e11), h11(e10))) | ~ (h11(op1(e10, e13)) = op2(h11(e10), h11(e13))) | ~ (h11(op1(e10, e12)) = op2(h11(e10), h11(e12))) | ~ (h11(op1(e10, e11)) = op2(h11(e10), h11(e11))) | ~ (h11(op1(e10, e10)) = op2(h11(e10), h11(e10)))), inference(subsumption_resolution, [], [f625, f432])).
fof(f625, plain, (~ (e23 = h11(e11)) | sP38 | sP37 | sP36 | ~ (h11(op1(e13, e13)) = op2(h11(e13), h11(e13))) | ~ (h11(op1(e13, e12)) = op2(h11(e13), h11(e12))) | ~ (h11(op1(e13, e11)) = op2(h11(e13), h11(e11))) | ~ (h11(op1(e13, e10)) = op2(h11(e13), h11(e10))) | ~ (h11(op1(e12, e13)) = op2(h11(e12), h11(e13))) | ~ (h11(op1(e12, e12)) = op2(h11(e12), h11(e12))) | ~ (h11(op1(e12, e11)) = op2(h11(e12), h11(e11))) | ~ (h11(op1(e12, e10)) = op2(h11(e12), h11(e10))) | ~ (h11(op1(e11, e13)) = op2(h11(e11), h11(e13))) | ~ (h11(op1(e11, e12)) = op2(h11(e11), h11(e12))) | ~ (h11(op1(e11, e11)) = op2(h11(e11), h11(e11))) | ~ (h11(op1(e11, e10)) = op2(h11(e11), h11(e10))) | ~ (h11(op1(e10, e13)) = op2(h11(e10), h11(e13))) | ~ (h11(op1(e10, e12)) = op2(h11(e10), h11(e12))) | ~ (h11(op1(e10, e11)) = op2(h11(e10), h11(e11))) | ~ (h11(op1(e10, e10)) = op2(h11(e10), h11(e10)))), inference(cnf_transformation, [], [f75])).
fof(f75, plain, (((~ (e23 = h12(e13)) & ~ (e23 = h12(e12)) & ~ (e23 = h12(e11)) & ~ (e23 = h12(e10))) | sP41 | sP40 | sP39 | ~ (h12(op1(e13, e13)) = op2(h12(e13), h12(e13))) | ~ (h12(op1(e13, e12)) = op2(h12(e13), h12(e12))) | ~ (h12(op1(e13, e11)) = op2(h12(e13), h12(e11))) | ~ (h12(op1(e13, e10)) = op2(h12(e13), h12(e10))) | ~ (h12(op1(e12, e13)) = op2(h12(e12), h12(e13))) | ~ (h12(op1(e12, e12)) = op2(h12(e12), h12(e12))) | ~ (h12(op1(e12, e11)) = op2(h12(e12), h12(e11))) | ~ (h12(op1(e12, e10)) = op2(h12(e12), h12(e10))) | ~ (h12(op1(e11, e13)) = op2(h12(e11), h12(e13))) | ~ (h12(op1(e11, e12)) = op2(h12(e11), h12(e12))) | ~ (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) | ~ (h12(op1(e11, e10)) = op2(h12(e11), h12(e10))) | ~ (h12(op1(e10, e13)) = op2(h12(e10), h12(e13))) | ~ (h12(op1(e10, e12)) = op2(h12(e10), h12(e12))) | ~ (h12(op1(e10, e11)) = op2(h12(e10), h12(e11))) | ~ (h12(op1(e10, e10)) = op2(h12(e10), h12(e10)))) & ((~ (e23 = h11(e13)) & ~ (e23 = h11(e12)) & ~ (e23 = h11(e11)) & ~ (e23 = h11(e10))) | sP38 | sP37 | sP36 | ~ (h11(op1(e13, e13)) = op2(h11(e13), h11(e13))) | ~ (h11(op1(e13, e12)) = op2(h11(e13), h11(e12))) | ~ (h11(op1(e13, e11)) = op2(h11(e13), h11(e11))) | ~ (h11(op1(e13, e10)) = op2(h11(e13), h11(e10))) | ~ (h11(op1(e12, e13)) = op2(h11(e12), h11(e13))) | ~ (h11(op1(e12, e12)) = op2(h11(e12), h11(e12))) | ~ (h11(op1(e12, e11)) = op2(h11(e12), h11(e11))) | ~ (h11(op1(e12, e10)) = op2(h11(e12), h11(e10))) | ~ (h11(op1(e11, e13)) = op2(h11(e11), h11(e13))) | ~ (h11(op1(e11, e12)) = op2(h11(e11), h11(e12))) | ~ (h11(op1(e11, e11)) = op2(h11(e11), h11(e11))) | ~ (h11(op1(e11, e10)) = op2(h11(e11), h11(e10))) | ~ (h11(op1(e10, e13)) = op2(h11(e10), h11(e13))) | ~ (h11(op1(e10, e12)) = op2(h11(e10), h11(e12))) | ~ (h11(op1(e10, e11)) = op2(h11(e10), h11(e11))) | ~ (h11(op1(e10, e10)) = op2(h11(e10), h11(e10)))) & ((~ (e23 = h10(e13)) & ~ (e23 = h10(e12)) & ~ (e23 = h10(e11)) & ~ (e23 = h10(e10))) | sP35 | sP34 | sP33 | ~ (h10(op1(e13, e13)) = op2(h10(e13), h10(e13))) | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) | ~ (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) | ~ (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) | ~ (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))) & ((~ (e23 = h9(e13)) & ~ (e23 = h9(e12)) & ~ (e23 = h9(e11)) & ~ (e23 = h9(e10))) | sP32 | sP31 | sP30 | ~ (h9(op1(e13, e13)) = op2(h9(e13), h9(e13))) | ~ (h9(op1(e13, e12)) = op2(h9(e13), h9(e12))) | ~ (h9(op1(e13, e11)) = op2(h9(e13), h9(e11))) | ~ (h9(op1(e13, e10)) = op2(h9(e13), h9(e10))) | ~ (h9(op1(e12, e13)) = op2(h9(e12), h9(e13))) | ~ (h9(op1(e12, e12)) = op2(h9(e12), h9(e12))) | ~ (h9(op1(e12, e11)) = op2(h9(e12), h9(e11))) | ~ (h9(op1(e12, e10)) = op2(h9(e12), h9(e10))) | ~ (h9(op1(e11, e13)) = op2(h9(e11), h9(e13))) | ~ (h9(op1(e11, e12)) = op2(h9(e11), h9(e12))) | ~ (h9(op1(e11, e11)) = op2(h9(e11), h9(e11))) | ~ (h9(op1(e11, e10)) = op2(h9(e11), h9(e10))) | ~ (h9(op1(e10, e13)) = op2(h9(e10), h9(e13))) | ~ (h9(op1(e10, e12)) = op2(h9(e10), h9(e12))) | ~ (h9(op1(e10, e11)) = op2(h9(e10), h9(e11))) | ~ (h9(op1(e10, e10)) = op2(h9(e10), h9(e10)))) & ((~ (e23 = h8(e13)) & ~ (e23 = h8(e12)) & ~ (e23 = h8(e11)) & ~ (e23 = h8(e10))) | sP29 | sP28 | sP27 | ~ (h8(op1(e13, e13)) = op2(h8(e13), h8(e13))) | ~ (h8(op1(e13, e12)) = op2(h8(e13), h8(e12))) | ~ (h8(op1(e13, e11)) = op2(h8(e13), h8(e11))) | ~ (h8(op1(e13, e10)) = op2(h8(e13), h8(e10))) | ~ (h8(op1(e12, e13)) = op2(h8(e12), h8(e13))) | ~ (h8(op1(e12, e12)) = op2(h8(e12), h8(e12))) | ~ (h8(op1(e12, e11)) = op2(h8(e12), h8(e11))) | ~ (h8(op1(e12, e10)) = op2(h8(e12), h8(e10))) | ~ (h8(op1(e11, e13)) = op2(h8(e11), h8(e13))) | ~ (h8(op1(e11, e12)) = op2(h8(e11), h8(e12))) | ~ (h8(op1(e11, e11)) = op2(h8(e11), h8(e11))) | ~ (h8(op1(e11, e10)) = op2(h8(e11), h8(e10))) | ~ (h8(op1(e10, e13)) = op2(h8(e10), h8(e13))) | ~ (h8(op1(e10, e12)) = op2(h8(e10), h8(e12))) | ~ (h8(op1(e10, e11)) = op2(h8(e10), h8(e11))) | ~ (h8(op1(e10, e10)) = op2(h8(e10), h8(e10)))) & ((~ (e23 = h7(e13)) & ~ (e23 = h7(e12)) & ~ (e23 = h7(e11)) & ~ (e23 = h7(e10))) | sP26 | sP25 | sP24 | ~ (h7(op1(e13, e13)) = op2(h7(e13), h7(e13))) | ~ (h7(op1(e13, e12)) = op2(h7(e13), h7(e12))) | ~ (h7(op1(e13, e11)) = op2(h7(e13), h7(e11))) | ~ (h7(op1(e13, e10)) = op2(h7(e13), h7(e10))) | ~ (h7(op1(e12, e13)) = op2(h7(e12), h7(e13))) | ~ (h7(op1(e12, e12)) = op2(h7(e12), h7(e12))) | ~ (h7(op1(e12, e11)) = op2(h7(e12), h7(e11))) | ~ (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) | ~ (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) | ~ (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) | ~ (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) | ~ (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) | ~ (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) | ~ (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) | ~ (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))) & ((~ (e23 = h6(e13)) & ~ (e23 = h6(e12)) & ~ (e23 = h6(e11)) & ~ (e23 = h6(e10))) | sP23 | sP22 | sP21 | ~ (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) | ~ (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))) & ((~ (e23 = h5(e13)) & ~ (e23 = h5(e12)) & ~ (e23 = h5(e11)) & ~ (e23 = h5(e10))) | sP20 | sP19 | sP18 | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))) & ((~ (e23 = h4(e13)) & ~ (e23 = h4(e12)) & ~ (e23 = h4(e11)) & ~ (e23 = h4(e10))) | sP17 | sP16 | sP15 | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) & ((~ (e23 = h3(e13)) & ~ (e23 = h3(e12)) & ~ (e23 = h3(e11)) & ~ (e23 = h3(e10))) | sP14 | sP13 | sP12 | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) & ((~ (e23 = h2(e13)) & ~ (e23 = h2(e12)) & ~ (e23 = h2(e11)) & ~ (e23 = h2(e10))) | sP11 | sP10 | sP9 | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) & ((~ (e23 = h1(e13)) & ~ (e23 = h1(e12)) & ~ (e23 = h1(e11)) & ~ (e23 = h1(e10))) | sP8 | sP7 | sP6 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(definition_folding, [], [f30, e74, e73, e72, e71, e70, e69, e68, e67, e66, e65, e64, e63, e62, e61, e60, e59, e58, e57, e56, e55, e54, e53, e52, e51, e50, e49, e48, e47, e46, e45, e44, e43, e42, e41, e40, e39])).
fof(f39, plain, ((~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10))) | ~ sP6), inference(usedef, [], [e39])).
fof(e39, plain, (sP6 <=> (~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f40, plain, ((~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10))) | ~ sP7), inference(usedef, [], [e40])).
fof(e40, plain, (sP7 <=> (~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f41, plain, ((~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10))) | ~ sP8), inference(usedef, [], [e41])).
fof(e41, plain, (sP8 <=> (~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f42, plain, ((~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10))) | ~ sP9), inference(usedef, [], [e42])).
fof(e42, plain, (sP9 <=> (~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f43, plain, ((~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10))) | ~ sP10), inference(usedef, [], [e43])).
fof(e43, plain, (sP10 <=> (~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f44, plain, ((~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10))) | ~ sP11), inference(usedef, [], [e44])).
fof(e44, plain, (sP11 <=> (~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f45, plain, ((~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10))) | ~ sP12), inference(usedef, [], [e45])).
fof(e45, plain, (sP12 <=> (~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f46, plain, ((~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10))) | ~ sP13), inference(usedef, [], [e46])).
fof(e46, plain, (sP13 <=> (~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f47, plain, ((~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10))) | ~ sP14), inference(usedef, [], [e47])).
fof(e47, plain, (sP14 <=> (~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f48, plain, ((~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10))) | ~ sP15), inference(usedef, [], [e48])).
fof(e48, plain, (sP15 <=> (~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f49, plain, ((~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10))) | ~ sP16), inference(usedef, [], [e49])).
fof(e49, plain, (sP16 <=> (~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f50, plain, ((~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10))) | ~ sP17), inference(usedef, [], [e50])).
fof(e50, plain, (sP17 <=> (~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f51, plain, ((~ (e20 = h5(e13)) & ~ (e20 = h5(e12)) & ~ (e20 = h5(e11)) & ~ (e20 = h5(e10))) | ~ sP18), inference(usedef, [], [e51])).
fof(e51, plain, (sP18 <=> (~ (e20 = h5(e13)) & ~ (e20 = h5(e12)) & ~ (e20 = h5(e11)) & ~ (e20 = h5(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f52, plain, ((~ (e21 = h5(e13)) & ~ (e21 = h5(e12)) & ~ (e21 = h5(e11)) & ~ (e21 = h5(e10))) | ~ sP19), inference(usedef, [], [e52])).
fof(e52, plain, (sP19 <=> (~ (e21 = h5(e13)) & ~ (e21 = h5(e12)) & ~ (e21 = h5(e11)) & ~ (e21 = h5(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP19])])).
fof(f53, plain, ((~ (e22 = h5(e13)) & ~ (e22 = h5(e12)) & ~ (e22 = h5(e11)) & ~ (e22 = h5(e10))) | ~ sP20), inference(usedef, [], [e53])).
fof(e53, plain, (sP20 <=> (~ (e22 = h5(e13)) & ~ (e22 = h5(e12)) & ~ (e22 = h5(e11)) & ~ (e22 = h5(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP20])])).
fof(f54, plain, ((~ (e20 = h6(e13)) & ~ (e20 = h6(e12)) & ~ (e20 = h6(e11)) & ~ (e20 = h6(e10))) | ~ sP21), inference(usedef, [], [e54])).
fof(e54, plain, (sP21 <=> (~ (e20 = h6(e13)) & ~ (e20 = h6(e12)) & ~ (e20 = h6(e11)) & ~ (e20 = h6(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP21])])).
fof(f55, plain, ((~ (e21 = h6(e13)) & ~ (e21 = h6(e12)) & ~ (e21 = h6(e11)) & ~ (e21 = h6(e10))) | ~ sP22), inference(usedef, [], [e55])).
fof(e55, plain, (sP22 <=> (~ (e21 = h6(e13)) & ~ (e21 = h6(e12)) & ~ (e21 = h6(e11)) & ~ (e21 = h6(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP22])])).
fof(f56, plain, ((~ (e22 = h6(e13)) & ~ (e22 = h6(e12)) & ~ (e22 = h6(e11)) & ~ (e22 = h6(e10))) | ~ sP23), inference(usedef, [], [e56])).
fof(e56, plain, (sP23 <=> (~ (e22 = h6(e13)) & ~ (e22 = h6(e12)) & ~ (e22 = h6(e11)) & ~ (e22 = h6(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP23])])).
fof(f57, plain, ((~ (e20 = h7(e13)) & ~ (e20 = h7(e12)) & ~ (e20 = h7(e11)) & ~ (e20 = h7(e10))) | ~ sP24), inference(usedef, [], [e57])).
fof(e57, plain, (sP24 <=> (~ (e20 = h7(e13)) & ~ (e20 = h7(e12)) & ~ (e20 = h7(e11)) & ~ (e20 = h7(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP24])])).
fof(f58, plain, ((~ (e21 = h7(e13)) & ~ (e21 = h7(e12)) & ~ (e21 = h7(e11)) & ~ (e21 = h7(e10))) | ~ sP25), inference(usedef, [], [e58])).
fof(e58, plain, (sP25 <=> (~ (e21 = h7(e13)) & ~ (e21 = h7(e12)) & ~ (e21 = h7(e11)) & ~ (e21 = h7(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP25])])).
fof(f59, plain, ((~ (e22 = h7(e13)) & ~ (e22 = h7(e12)) & ~ (e22 = h7(e11)) & ~ (e22 = h7(e10))) | ~ sP26), inference(usedef, [], [e59])).
fof(e59, plain, (sP26 <=> (~ (e22 = h7(e13)) & ~ (e22 = h7(e12)) & ~ (e22 = h7(e11)) & ~ (e22 = h7(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP26])])).
fof(f60, plain, ((~ (e20 = h8(e13)) & ~ (e20 = h8(e12)) & ~ (e20 = h8(e11)) & ~ (e20 = h8(e10))) | ~ sP27), inference(usedef, [], [e60])).
fof(e60, plain, (sP27 <=> (~ (e20 = h8(e13)) & ~ (e20 = h8(e12)) & ~ (e20 = h8(e11)) & ~ (e20 = h8(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP27])])).
fof(f61, plain, ((~ (e21 = h8(e13)) & ~ (e21 = h8(e12)) & ~ (e21 = h8(e11)) & ~ (e21 = h8(e10))) | ~ sP28), inference(usedef, [], [e61])).
fof(e61, plain, (sP28 <=> (~ (e21 = h8(e13)) & ~ (e21 = h8(e12)) & ~ (e21 = h8(e11)) & ~ (e21 = h8(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP28])])).
fof(f62, plain, ((~ (e22 = h8(e13)) & ~ (e22 = h8(e12)) & ~ (e22 = h8(e11)) & ~ (e22 = h8(e10))) | ~ sP29), inference(usedef, [], [e62])).
fof(e62, plain, (sP29 <=> (~ (e22 = h8(e13)) & ~ (e22 = h8(e12)) & ~ (e22 = h8(e11)) & ~ (e22 = h8(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP29])])).
fof(f63, plain, ((~ (e20 = h9(e13)) & ~ (e20 = h9(e12)) & ~ (e20 = h9(e11)) & ~ (e20 = h9(e10))) | ~ sP30), inference(usedef, [], [e63])).
fof(e63, plain, (sP30 <=> (~ (e20 = h9(e13)) & ~ (e20 = h9(e12)) & ~ (e20 = h9(e11)) & ~ (e20 = h9(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP30])])).
fof(f64, plain, ((~ (e21 = h9(e13)) & ~ (e21 = h9(e12)) & ~ (e21 = h9(e11)) & ~ (e21 = h9(e10))) | ~ sP31), inference(usedef, [], [e64])).
fof(e64, plain, (sP31 <=> (~ (e21 = h9(e13)) & ~ (e21 = h9(e12)) & ~ (e21 = h9(e11)) & ~ (e21 = h9(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP31])])).
fof(f65, plain, ((~ (e22 = h9(e13)) & ~ (e22 = h9(e12)) & ~ (e22 = h9(e11)) & ~ (e22 = h9(e10))) | ~ sP32), inference(usedef, [], [e65])).
fof(e65, plain, (sP32 <=> (~ (e22 = h9(e13)) & ~ (e22 = h9(e12)) & ~ (e22 = h9(e11)) & ~ (e22 = h9(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP32])])).
fof(f66, plain, ((~ (e20 = h10(e13)) & ~ (e20 = h10(e12)) & ~ (e20 = h10(e11)) & ~ (e20 = h10(e10))) | ~ sP33), inference(usedef, [], [e66])).
fof(e66, plain, (sP33 <=> (~ (e20 = h10(e13)) & ~ (e20 = h10(e12)) & ~ (e20 = h10(e11)) & ~ (e20 = h10(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP33])])).
fof(f67, plain, ((~ (e21 = h10(e13)) & ~ (e21 = h10(e12)) & ~ (e21 = h10(e11)) & ~ (e21 = h10(e10))) | ~ sP34), inference(usedef, [], [e67])).
fof(e67, plain, (sP34 <=> (~ (e21 = h10(e13)) & ~ (e21 = h10(e12)) & ~ (e21 = h10(e11)) & ~ (e21 = h10(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP34])])).
fof(f68, plain, ((~ (e22 = h10(e13)) & ~ (e22 = h10(e12)) & ~ (e22 = h10(e11)) & ~ (e22 = h10(e10))) | ~ sP35), inference(usedef, [], [e68])).
fof(e68, plain, (sP35 <=> (~ (e22 = h10(e13)) & ~ (e22 = h10(e12)) & ~ (e22 = h10(e11)) & ~ (e22 = h10(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP35])])).
fof(f69, plain, ((~ (e20 = h11(e13)) & ~ (e20 = h11(e12)) & ~ (e20 = h11(e11)) & ~ (e20 = h11(e10))) | ~ sP36), inference(usedef, [], [e69])).
fof(e69, plain, (sP36 <=> (~ (e20 = h11(e13)) & ~ (e20 = h11(e12)) & ~ (e20 = h11(e11)) & ~ (e20 = h11(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP36])])).
fof(f71, plain, ((~ (e22 = h11(e13)) & ~ (e22 = h11(e12)) & ~ (e22 = h11(e11)) & ~ (e22 = h11(e10))) | ~ sP38), inference(usedef, [], [e71])).
fof(e71, plain, (sP38 <=> (~ (e22 = h11(e13)) & ~ (e22 = h11(e12)) & ~ (e22 = h11(e11)) & ~ (e22 = h11(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP38])])).
fof(f72, plain, ((~ (e20 = h12(e13)) & ~ (e20 = h12(e12)) & ~ (e20 = h12(e11)) & ~ (e20 = h12(e10))) | ~ sP39), inference(usedef, [], [e72])).
fof(e72, plain, (sP39 <=> (~ (e20 = h12(e13)) & ~ (e20 = h12(e12)) & ~ (e20 = h12(e11)) & ~ (e20 = h12(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP39])])).
fof(f73, plain, ((~ (e21 = h12(e13)) & ~ (e21 = h12(e12)) & ~ (e21 = h12(e11)) & ~ (e21 = h12(e10))) | ~ sP40), inference(usedef, [], [e73])).
fof(e73, plain, (sP40 <=> (~ (e21 = h12(e13)) & ~ (e21 = h12(e12)) & ~ (e21 = h12(e11)) & ~ (e21 = h12(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP40])])).
fof(f74, plain, ((~ (e22 = h12(e13)) & ~ (e22 = h12(e12)) & ~ (e22 = h12(e11)) & ~ (e22 = h12(e10))) | ~ sP41), inference(usedef, [], [e74])).
fof(e74, plain, (sP41 <=> (~ (e22 = h12(e13)) & ~ (e22 = h12(e12)) & ~ (e22 = h12(e11)) & ~ (e22 = h12(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP41])])).
fof(f30, plain, (((~ (e23 = h12(e13)) & ~ (e23 = h12(e12)) & ~ (e23 = h12(e11)) & ~ (e23 = h12(e10))) | (~ (e22 = h12(e13)) & ~ (e22 = h12(e12)) & ~ (e22 = h12(e11)) & ~ (e22 = h12(e10))) | (~ (e21 = h12(e13)) & ~ (e21 = h12(e12)) & ~ (e21 = h12(e11)) & ~ (e21 = h12(e10))) | (~ (e20 = h12(e13)) & ~ (e20 = h12(e12)) & ~ (e20 = h12(e11)) & ~ (e20 = h12(e10))) | ~ (h12(op1(e13, e13)) = op2(h12(e13), h12(e13))) | ~ (h12(op1(e13, e12)) = op2(h12(e13), h12(e12))) | ~ (h12(op1(e13, e11)) = op2(h12(e13), h12(e11))) | ~ (h12(op1(e13, e10)) = op2(h12(e13), h12(e10))) | ~ (h12(op1(e12, e13)) = op2(h12(e12), h12(e13))) | ~ (h12(op1(e12, e12)) = op2(h12(e12), h12(e12))) | ~ (h12(op1(e12, e11)) = op2(h12(e12), h12(e11))) | ~ (h12(op1(e12, e10)) = op2(h12(e12), h12(e10))) | ~ (h12(op1(e11, e13)) = op2(h12(e11), h12(e13))) | ~ (h12(op1(e11, e12)) = op2(h12(e11), h12(e12))) | ~ (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) | ~ (h12(op1(e11, e10)) = op2(h12(e11), h12(e10))) | ~ (h12(op1(e10, e13)) = op2(h12(e10), h12(e13))) | ~ (h12(op1(e10, e12)) = op2(h12(e10), h12(e12))) | ~ (h12(op1(e10, e11)) = op2(h12(e10), h12(e11))) | ~ (h12(op1(e10, e10)) = op2(h12(e10), h12(e10)))) & ((~ (e23 = h11(e13)) & ~ (e23 = h11(e12)) & ~ (e23 = h11(e11)) & ~ (e23 = h11(e10))) | (~ (e22 = h11(e13)) & ~ (e22 = h11(e12)) & ~ (e22 = h11(e11)) & ~ (e22 = h11(e10))) | (~ (e21 = h11(e13)) & ~ (e21 = h11(e12)) & ~ (e21 = h11(e11)) & ~ (e21 = h11(e10))) | (~ (e20 = h11(e13)) & ~ (e20 = h11(e12)) & ~ (e20 = h11(e11)) & ~ (e20 = h11(e10))) | ~ (h11(op1(e13, e13)) = op2(h11(e13), h11(e13))) | ~ (h11(op1(e13, e12)) = op2(h11(e13), h11(e12))) | ~ (h11(op1(e13, e11)) = op2(h11(e13), h11(e11))) | ~ (h11(op1(e13, e10)) = op2(h11(e13), h11(e10))) | ~ (h11(op1(e12, e13)) = op2(h11(e12), h11(e13))) | ~ (h11(op1(e12, e12)) = op2(h11(e12), h11(e12))) | ~ (h11(op1(e12, e11)) = op2(h11(e12), h11(e11))) | ~ (h11(op1(e12, e10)) = op2(h11(e12), h11(e10))) | ~ (h11(op1(e11, e13)) = op2(h11(e11), h11(e13))) | ~ (h11(op1(e11, e12)) = op2(h11(e11), h11(e12))) | ~ (h11(op1(e11, e11)) = op2(h11(e11), h11(e11))) | ~ (h11(op1(e11, e10)) = op2(h11(e11), h11(e10))) | ~ (h11(op1(e10, e13)) = op2(h11(e10), h11(e13))) | ~ (h11(op1(e10, e12)) = op2(h11(e10), h11(e12))) | ~ (h11(op1(e10, e11)) = op2(h11(e10), h11(e11))) | ~ (h11(op1(e10, e10)) = op2(h11(e10), h11(e10)))) & ((~ (e23 = h10(e13)) & ~ (e23 = h10(e12)) & ~ (e23 = h10(e11)) & ~ (e23 = h10(e10))) | (~ (e22 = h10(e13)) & ~ (e22 = h10(e12)) & ~ (e22 = h10(e11)) & ~ (e22 = h10(e10))) | (~ (e21 = h10(e13)) & ~ (e21 = h10(e12)) & ~ (e21 = h10(e11)) & ~ (e21 = h10(e10))) | (~ (e20 = h10(e13)) & ~ (e20 = h10(e12)) & ~ (e20 = h10(e11)) & ~ (e20 = h10(e10))) | ~ (h10(op1(e13, e13)) = op2(h10(e13), h10(e13))) | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) | ~ (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) | ~ (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) | ~ (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))) & ((~ (e23 = h9(e13)) & ~ (e23 = h9(e12)) & ~ (e23 = h9(e11)) & ~ (e23 = h9(e10))) | (~ (e22 = h9(e13)) & ~ (e22 = h9(e12)) & ~ (e22 = h9(e11)) & ~ (e22 = h9(e10))) | (~ (e21 = h9(e13)) & ~ (e21 = h9(e12)) & ~ (e21 = h9(e11)) & ~ (e21 = h9(e10))) | (~ (e20 = h9(e13)) & ~ (e20 = h9(e12)) & ~ (e20 = h9(e11)) & ~ (e20 = h9(e10))) | ~ (h9(op1(e13, e13)) = op2(h9(e13), h9(e13))) | ~ (h9(op1(e13, e12)) = op2(h9(e13), h9(e12))) | ~ (h9(op1(e13, e11)) = op2(h9(e13), h9(e11))) | ~ (h9(op1(e13, e10)) = op2(h9(e13), h9(e10))) | ~ (h9(op1(e12, e13)) = op2(h9(e12), h9(e13))) | ~ (h9(op1(e12, e12)) = op2(h9(e12), h9(e12))) | ~ (h9(op1(e12, e11)) = op2(h9(e12), h9(e11))) | ~ (h9(op1(e12, e10)) = op2(h9(e12), h9(e10))) | ~ (h9(op1(e11, e13)) = op2(h9(e11), h9(e13))) | ~ (h9(op1(e11, e12)) = op2(h9(e11), h9(e12))) | ~ (h9(op1(e11, e11)) = op2(h9(e11), h9(e11))) | ~ (h9(op1(e11, e10)) = op2(h9(e11), h9(e10))) | ~ (h9(op1(e10, e13)) = op2(h9(e10), h9(e13))) | ~ (h9(op1(e10, e12)) = op2(h9(e10), h9(e12))) | ~ (h9(op1(e10, e11)) = op2(h9(e10), h9(e11))) | ~ (h9(op1(e10, e10)) = op2(h9(e10), h9(e10)))) & ((~ (e23 = h8(e13)) & ~ (e23 = h8(e12)) & ~ (e23 = h8(e11)) & ~ (e23 = h8(e10))) | (~ (e22 = h8(e13)) & ~ (e22 = h8(e12)) & ~ (e22 = h8(e11)) & ~ (e22 = h8(e10))) | (~ (e21 = h8(e13)) & ~ (e21 = h8(e12)) & ~ (e21 = h8(e11)) & ~ (e21 = h8(e10))) | (~ (e20 = h8(e13)) & ~ (e20 = h8(e12)) & ~ (e20 = h8(e11)) & ~ (e20 = h8(e10))) | ~ (h8(op1(e13, e13)) = op2(h8(e13), h8(e13))) | ~ (h8(op1(e13, e12)) = op2(h8(e13), h8(e12))) | ~ (h8(op1(e13, e11)) = op2(h8(e13), h8(e11))) | ~ (h8(op1(e13, e10)) = op2(h8(e13), h8(e10))) | ~ (h8(op1(e12, e13)) = op2(h8(e12), h8(e13))) | ~ (h8(op1(e12, e12)) = op2(h8(e12), h8(e12))) | ~ (h8(op1(e12, e11)) = op2(h8(e12), h8(e11))) | ~ (h8(op1(e12, e10)) = op2(h8(e12), h8(e10))) | ~ (h8(op1(e11, e13)) = op2(h8(e11), h8(e13))) | ~ (h8(op1(e11, e12)) = op2(h8(e11), h8(e12))) | ~ (h8(op1(e11, e11)) = op2(h8(e11), h8(e11))) | ~ (h8(op1(e11, e10)) = op2(h8(e11), h8(e10))) | ~ (h8(op1(e10, e13)) = op2(h8(e10), h8(e13))) | ~ (h8(op1(e10, e12)) = op2(h8(e10), h8(e12))) | ~ (h8(op1(e10, e11)) = op2(h8(e10), h8(e11))) | ~ (h8(op1(e10, e10)) = op2(h8(e10), h8(e10)))) & ((~ (e23 = h7(e13)) & ~ (e23 = h7(e12)) & ~ (e23 = h7(e11)) & ~ (e23 = h7(e10))) | (~ (e22 = h7(e13)) & ~ (e22 = h7(e12)) & ~ (e22 = h7(e11)) & ~ (e22 = h7(e10))) | (~ (e21 = h7(e13)) & ~ (e21 = h7(e12)) & ~ (e21 = h7(e11)) & ~ (e21 = h7(e10))) | (~ (e20 = h7(e13)) & ~ (e20 = h7(e12)) & ~ (e20 = h7(e11)) & ~ (e20 = h7(e10))) | ~ (h7(op1(e13, e13)) = op2(h7(e13), h7(e13))) | ~ (h7(op1(e13, e12)) = op2(h7(e13), h7(e12))) | ~ (h7(op1(e13, e11)) = op2(h7(e13), h7(e11))) | ~ (h7(op1(e13, e10)) = op2(h7(e13), h7(e10))) | ~ (h7(op1(e12, e13)) = op2(h7(e12), h7(e13))) | ~ (h7(op1(e12, e12)) = op2(h7(e12), h7(e12))) | ~ (h7(op1(e12, e11)) = op2(h7(e12), h7(e11))) | ~ (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) | ~ (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) | ~ (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) | ~ (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) | ~ (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) | ~ (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) | ~ (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) | ~ (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))) & ((~ (e23 = h6(e13)) & ~ (e23 = h6(e12)) & ~ (e23 = h6(e11)) & ~ (e23 = h6(e10))) | (~ (e22 = h6(e13)) & ~ (e22 = h6(e12)) & ~ (e22 = h6(e11)) & ~ (e22 = h6(e10))) | (~ (e21 = h6(e13)) & ~ (e21 = h6(e12)) & ~ (e21 = h6(e11)) & ~ (e21 = h6(e10))) | (~ (e20 = h6(e13)) & ~ (e20 = h6(e12)) & ~ (e20 = h6(e11)) & ~ (e20 = h6(e10))) | ~ (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) | ~ (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))) & ((~ (e23 = h5(e13)) & ~ (e23 = h5(e12)) & ~ (e23 = h5(e11)) & ~ (e23 = h5(e10))) | (~ (e22 = h5(e13)) & ~ (e22 = h5(e12)) & ~ (e22 = h5(e11)) & ~ (e22 = h5(e10))) | (~ (e21 = h5(e13)) & ~ (e21 = h5(e12)) & ~ (e21 = h5(e11)) & ~ (e21 = h5(e10))) | (~ (e20 = h5(e13)) & ~ (e20 = h5(e12)) & ~ (e20 = h5(e11)) & ~ (e20 = h5(e10))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))) & ((~ (e23 = h4(e13)) & ~ (e23 = h4(e12)) & ~ (e23 = h4(e11)) & ~ (e23 = h4(e10))) | (~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10))) | (~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10))) | (~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10))) | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) & ((~ (e23 = h3(e13)) & ~ (e23 = h3(e12)) & ~ (e23 = h3(e11)) & ~ (e23 = h3(e10))) | (~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10))) | (~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10))) | (~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10))) | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) & ((~ (e23 = h2(e13)) & ~ (e23 = h2(e12)) & ~ (e23 = h2(e11)) & ~ (e23 = h2(e10))) | (~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10))) | (~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10))) | (~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) & ((~ (e23 = h1(e13)) & ~ (e23 = h1(e12)) & ~ (e23 = h1(e11)) & ~ (e23 = h1(e10))) | (~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10))) | (~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10))) | (~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10))) | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(ennf_transformation, [], [f29])).
fof(f29, plain, ~ ((((e23 = h12(e13)) | (e23 = h12(e12)) | (e23 = h12(e11)) | (e23 = h12(e10))) & ((e22 = h12(e13)) | (e22 = h12(e12)) | (e22 = h12(e11)) | (e22 = h12(e10))) & ((e21 = h12(e13)) | (e21 = h12(e12)) | (e21 = h12(e11)) | (e21 = h12(e10))) & ((e20 = h12(e13)) | (e20 = h12(e12)) | (e20 = h12(e11)) | (e20 = h12(e10))) & (h12(op1(e13, e13)) = op2(h12(e13), h12(e13))) & (h12(op1(e13, e12)) = op2(h12(e13), h12(e12))) & (h12(op1(e13, e11)) = op2(h12(e13), h12(e11))) & (h12(op1(e13, e10)) = op2(h12(e13), h12(e10))) & (h12(op1(e12, e13)) = op2(h12(e12), h12(e13))) & (h12(op1(e12, e12)) = op2(h12(e12), h12(e12))) & (h12(op1(e12, e11)) = op2(h12(e12), h12(e11))) & (h12(op1(e12, e10)) = op2(h12(e12), h12(e10))) & (h12(op1(e11, e13)) = op2(h12(e11), h12(e13))) & (h12(op1(e11, e12)) = op2(h12(e11), h12(e12))) & (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) & (h12(op1(e11, e10)) = op2(h12(e11), h12(e10))) & (h12(op1(e10, e13)) = op2(h12(e10), h12(e13))) & (h12(op1(e10, e12)) = op2(h12(e10), h12(e12))) & (h12(op1(e10, e11)) = op2(h12(e10), h12(e11))) & (h12(op1(e10, e10)) = op2(h12(e10), h12(e10)))) | (((e23 = h11(e13)) | (e23 = h11(e12)) | (e23 = h11(e11)) | (e23 = h11(e10))) & ((e22 = h11(e13)) | (e22 = h11(e12)) | (e22 = h11(e11)) | (e22 = h11(e10))) & ((e21 = h11(e13)) | (e21 = h11(e12)) | (e21 = h11(e11)) | (e21 = h11(e10))) & ((e20 = h11(e13)) | (e20 = h11(e12)) | (e20 = h11(e11)) | (e20 = h11(e10))) & (h11(op1(e13, e13)) = op2(h11(e13), h11(e13))) & (h11(op1(e13, e12)) = op2(h11(e13), h11(e12))) & (h11(op1(e13, e11)) = op2(h11(e13), h11(e11))) & (h11(op1(e13, e10)) = op2(h11(e13), h11(e10))) & (h11(op1(e12, e13)) = op2(h11(e12), h11(e13))) & (h11(op1(e12, e12)) = op2(h11(e12), h11(e12))) & (h11(op1(e12, e11)) = op2(h11(e12), h11(e11))) & (h11(op1(e12, e10)) = op2(h11(e12), h11(e10))) & (h11(op1(e11, e13)) = op2(h11(e11), h11(e13))) & (h11(op1(e11, e12)) = op2(h11(e11), h11(e12))) & (h11(op1(e11, e11)) = op2(h11(e11), h11(e11))) & (h11(op1(e11, e10)) = op2(h11(e11), h11(e10))) & (h11(op1(e10, e13)) = op2(h11(e10), h11(e13))) & (h11(op1(e10, e12)) = op2(h11(e10), h11(e12))) & (h11(op1(e10, e11)) = op2(h11(e10), h11(e11))) & (h11(op1(e10, e10)) = op2(h11(e10), h11(e10)))) | (((e23 = h10(e13)) | (e23 = h10(e12)) | (e23 = h10(e11)) | (e23 = h10(e10))) & ((e22 = h10(e13)) | (e22 = h10(e12)) | (e22 = h10(e11)) | (e22 = h10(e10))) & ((e21 = h10(e13)) | (e21 = h10(e12)) | (e21 = h10(e11)) | (e21 = h10(e10))) & ((e20 = h10(e13)) | (e20 = h10(e12)) | (e20 = h10(e11)) | (e20 = h10(e10))) & (h10(op1(e13, e13)) = op2(h10(e13), h10(e13))) & (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) & (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) & (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) & (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) & (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) & (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) & (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) & (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) & (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) & (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) & (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) & (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) & (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) & (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) & (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))) | (((e23 = h9(e13)) | (e23 = h9(e12)) | (e23 = h9(e11)) | (e23 = h9(e10))) & ((e22 = h9(e13)) | (e22 = h9(e12)) | (e22 = h9(e11)) | (e22 = h9(e10))) & ((e21 = h9(e13)) | (e21 = h9(e12)) | (e21 = h9(e11)) | (e21 = h9(e10))) & ((e20 = h9(e13)) | (e20 = h9(e12)) | (e20 = h9(e11)) | (e20 = h9(e10))) & (h9(op1(e13, e13)) = op2(h9(e13), h9(e13))) & (h9(op1(e13, e12)) = op2(h9(e13), h9(e12))) & (h9(op1(e13, e11)) = op2(h9(e13), h9(e11))) & (h9(op1(e13, e10)) = op2(h9(e13), h9(e10))) & (h9(op1(e12, e13)) = op2(h9(e12), h9(e13))) & (h9(op1(e12, e12)) = op2(h9(e12), h9(e12))) & (h9(op1(e12, e11)) = op2(h9(e12), h9(e11))) & (h9(op1(e12, e10)) = op2(h9(e12), h9(e10))) & (h9(op1(e11, e13)) = op2(h9(e11), h9(e13))) & (h9(op1(e11, e12)) = op2(h9(e11), h9(e12))) & (h9(op1(e11, e11)) = op2(h9(e11), h9(e11))) & (h9(op1(e11, e10)) = op2(h9(e11), h9(e10))) & (h9(op1(e10, e13)) = op2(h9(e10), h9(e13))) & (h9(op1(e10, e12)) = op2(h9(e10), h9(e12))) & (h9(op1(e10, e11)) = op2(h9(e10), h9(e11))) & (h9(op1(e10, e10)) = op2(h9(e10), h9(e10)))) | (((e23 = h8(e13)) | (e23 = h8(e12)) | (e23 = h8(e11)) | (e23 = h8(e10))) & ((e22 = h8(e13)) | (e22 = h8(e12)) | (e22 = h8(e11)) | (e22 = h8(e10))) & ((e21 = h8(e13)) | (e21 = h8(e12)) | (e21 = h8(e11)) | (e21 = h8(e10))) & ((e20 = h8(e13)) | (e20 = h8(e12)) | (e20 = h8(e11)) | (e20 = h8(e10))) & (h8(op1(e13, e13)) = op2(h8(e13), h8(e13))) & (h8(op1(e13, e12)) = op2(h8(e13), h8(e12))) & (h8(op1(e13, e11)) = op2(h8(e13), h8(e11))) & (h8(op1(e13, e10)) = op2(h8(e13), h8(e10))) & (h8(op1(e12, e13)) = op2(h8(e12), h8(e13))) & (h8(op1(e12, e12)) = op2(h8(e12), h8(e12))) & (h8(op1(e12, e11)) = op2(h8(e12), h8(e11))) & (h8(op1(e12, e10)) = op2(h8(e12), h8(e10))) & (h8(op1(e11, e13)) = op2(h8(e11), h8(e13))) & (h8(op1(e11, e12)) = op2(h8(e11), h8(e12))) & (h8(op1(e11, e11)) = op2(h8(e11), h8(e11))) & (h8(op1(e11, e10)) = op2(h8(e11), h8(e10))) & (h8(op1(e10, e13)) = op2(h8(e10), h8(e13))) & (h8(op1(e10, e12)) = op2(h8(e10), h8(e12))) & (h8(op1(e10, e11)) = op2(h8(e10), h8(e11))) & (h8(op1(e10, e10)) = op2(h8(e10), h8(e10)))) | (((e23 = h7(e13)) | (e23 = h7(e12)) | (e23 = h7(e11)) | (e23 = h7(e10))) & ((e22 = h7(e13)) | (e22 = h7(e12)) | (e22 = h7(e11)) | (e22 = h7(e10))) & ((e21 = h7(e13)) | (e21 = h7(e12)) | (e21 = h7(e11)) | (e21 = h7(e10))) & ((e20 = h7(e13)) | (e20 = h7(e12)) | (e20 = h7(e11)) | (e20 = h7(e10))) & (h7(op1(e13, e13)) = op2(h7(e13), h7(e13))) & (h7(op1(e13, e12)) = op2(h7(e13), h7(e12))) & (h7(op1(e13, e11)) = op2(h7(e13), h7(e11))) & (h7(op1(e13, e10)) = op2(h7(e13), h7(e10))) & (h7(op1(e12, e13)) = op2(h7(e12), h7(e13))) & (h7(op1(e12, e12)) = op2(h7(e12), h7(e12))) & (h7(op1(e12, e11)) = op2(h7(e12), h7(e11))) & (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) & (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) & (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) & (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) & (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) & (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) & (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) & (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) & (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))) | (((e23 = h6(e13)) | (e23 = h6(e12)) | (e23 = h6(e11)) | (e23 = h6(e10))) & ((e22 = h6(e13)) | (e22 = h6(e12)) | (e22 = h6(e11)) | (e22 = h6(e10))) & ((e21 = h6(e13)) | (e21 = h6(e12)) | (e21 = h6(e11)) | (e21 = h6(e10))) & ((e20 = h6(e13)) | (e20 = h6(e12)) | (e20 = h6(e11)) | (e20 = h6(e10))) & (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) & (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) & (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) & (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) & (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) & (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) & (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) & (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) & (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) & (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) & (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) & (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) & (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) & (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) & (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) & (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))) | (((e23 = h5(e13)) | (e23 = h5(e12)) | (e23 = h5(e11)) | (e23 = h5(e10))) & ((e22 = h5(e13)) | (e22 = h5(e12)) | (e22 = h5(e11)) | (e22 = h5(e10))) & ((e21 = h5(e13)) | (e21 = h5(e12)) | (e21 = h5(e11)) | (e21 = h5(e10))) & ((e20 = h5(e13)) | (e20 = h5(e12)) | (e20 = h5(e11)) | (e20 = h5(e10))) & (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) & (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) & (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) & (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) & (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) & (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) & (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) & (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) & (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) & (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) & (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) & (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) & (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) & (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) & (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) & (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))) | (((e23 = h4(e13)) | (e23 = h4(e12)) | (e23 = h4(e11)) | (e23 = h4(e10))) & ((e22 = h4(e13)) | (e22 = h4(e12)) | (e22 = h4(e11)) | (e22 = h4(e10))) & ((e21 = h4(e13)) | (e21 = h4(e12)) | (e21 = h4(e11)) | (e21 = h4(e10))) & ((e20 = h4(e13)) | (e20 = h4(e12)) | (e20 = h4(e11)) | (e20 = h4(e10))) & (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) & (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) & (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) & (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) & (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) & (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) & (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) & (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) & (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) & (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) & (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) & (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) & (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) & (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) & (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) & (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) | (((e23 = h3(e13)) | (e23 = h3(e12)) | (e23 = h3(e11)) | (e23 = h3(e10))) & ((e22 = h3(e13)) | (e22 = h3(e12)) | (e22 = h3(e11)) | (e22 = h3(e10))) & ((e21 = h3(e13)) | (e21 = h3(e12)) | (e21 = h3(e11)) | (e21 = h3(e10))) & ((e20 = h3(e13)) | (e20 = h3(e12)) | (e20 = h3(e11)) | (e20 = h3(e10))) & (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) & (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) & (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) & (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) & (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) & (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) & (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) & (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) & (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) & (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) & (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) & (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) & (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) & (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) & (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) & (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) | (((e23 = h2(e13)) | (e23 = h2(e12)) | (e23 = h2(e11)) | (e23 = h2(e10))) & ((e22 = h2(e13)) | (e22 = h2(e12)) | (e22 = h2(e11)) | (e22 = h2(e10))) & ((e21 = h2(e13)) | (e21 = h2(e12)) | (e21 = h2(e11)) | (e21 = h2(e10))) & ((e20 = h2(e13)) | (e20 = h2(e12)) | (e20 = h2(e11)) | (e20 = h2(e10))) & (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) & (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) & (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) & (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) & (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) & (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) & (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) & (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) & (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) & (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) & (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) & (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) & (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) & (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) & (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) & (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) | (((e23 = h1(e13)) | (e23 = h1(e12)) | (e23 = h1(e11)) | (e23 = h1(e10))) & ((e22 = h1(e13)) | (e22 = h1(e12)) | (e22 = h1(e11)) | (e22 = h1(e10))) & ((e21 = h1(e13)) | (e21 = h1(e12)) | (e21 = h1(e11)) | (e21 = h1(e10))) & ((e20 = h1(e13)) | (e20 = h1(e12)) | (e20 = h1(e11)) | (e20 = h1(e10))) & (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) & (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) & (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) & (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) & (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) & (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) & (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) & (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) & (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) & (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) & (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) & (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) & (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) & (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) & (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) & (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(negated_conjecture, [], [f28])).
fof(f28, plain, ~ ((((e23 = h12(e13)) | (e23 = h12(e12)) | (e23 = h12(e11)) | (e23 = h12(e10))) & ((e22 = h12(e13)) | (e22 = h12(e12)) | (e22 = h12(e11)) | (e22 = h12(e10))) & ((e21 = h12(e13)) | (e21 = h12(e12)) | (e21 = h12(e11)) | (e21 = h12(e10))) & ((e20 = h12(e13)) | (e20 = h12(e12)) | (e20 = h12(e11)) | (e20 = h12(e10))) & (h12(op1(e13, e13)) = op2(h12(e13), h12(e13))) & (h12(op1(e13, e12)) = op2(h12(e13), h12(e12))) & (h12(op1(e13, e11)) = op2(h12(e13), h12(e11))) & (h12(op1(e13, e10)) = op2(h12(e13), h12(e10))) & (h12(op1(e12, e13)) = op2(h12(e12), h12(e13))) & (h12(op1(e12, e12)) = op2(h12(e12), h12(e12))) & (h12(op1(e12, e11)) = op2(h12(e12), h12(e11))) & (h12(op1(e12, e10)) = op2(h12(e12), h12(e10))) & (h12(op1(e11, e13)) = op2(h12(e11), h12(e13))) & (h12(op1(e11, e12)) = op2(h12(e11), h12(e12))) & (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) & (h12(op1(e11, e10)) = op2(h12(e11), h12(e10))) & (h12(op1(e10, e13)) = op2(h12(e10), h12(e13))) & (h12(op1(e10, e12)) = op2(h12(e10), h12(e12))) & (h12(op1(e10, e11)) = op2(h12(e10), h12(e11))) & (h12(op1(e10, e10)) = op2(h12(e10), h12(e10)))) | (((e23 = h11(e13)) | (e23 = h11(e12)) | (e23 = h11(e11)) | (e23 = h11(e10))) & ((e22 = h11(e13)) | (e22 = h11(e12)) | (e22 = h11(e11)) | (e22 = h11(e10))) & ((e21 = h11(e13)) | (e21 = h11(e12)) | (e21 = h11(e11)) | (e21 = h11(e10))) & ((e20 = h11(e13)) | (e20 = h11(e12)) | (e20 = h11(e11)) | (e20 = h11(e10))) & (h11(op1(e13, e13)) = op2(h11(e13), h11(e13))) & (h11(op1(e13, e12)) = op2(h11(e13), h11(e12))) & (h11(op1(e13, e11)) = op2(h11(e13), h11(e11))) & (h11(op1(e13, e10)) = op2(h11(e13), h11(e10))) & (h11(op1(e12, e13)) = op2(h11(e12), h11(e13))) & (h11(op1(e12, e12)) = op2(h11(e12), h11(e12))) & (h11(op1(e12, e11)) = op2(h11(e12), h11(e11))) & (h11(op1(e12, e10)) = op2(h11(e12), h11(e10))) & (h11(op1(e11, e13)) = op2(h11(e11), h11(e13))) & (h11(op1(e11, e12)) = op2(h11(e11), h11(e12))) & (h11(op1(e11, e11)) = op2(h11(e11), h11(e11))) & (h11(op1(e11, e10)) = op2(h11(e11), h11(e10))) & (h11(op1(e10, e13)) = op2(h11(e10), h11(e13))) & (h11(op1(e10, e12)) = op2(h11(e10), h11(e12))) & (h11(op1(e10, e11)) = op2(h11(e10), h11(e11))) & (h11(op1(e10, e10)) = op2(h11(e10), h11(e10)))) | (((e23 = h10(e13)) | (e23 = h10(e12)) | (e23 = h10(e11)) | (e23 = h10(e10))) & ((e22 = h10(e13)) | (e22 = h10(e12)) | (e22 = h10(e11)) | (e22 = h10(e10))) & ((e21 = h10(e13)) | (e21 = h10(e12)) | (e21 = h10(e11)) | (e21 = h10(e10))) & ((e20 = h10(e13)) | (e20 = h10(e12)) | (e20 = h10(e11)) | (e20 = h10(e10))) & (h10(op1(e13, e13)) = op2(h10(e13), h10(e13))) & (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) & (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) & (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) & (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) & (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) & (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) & (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) & (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) & (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) & (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) & (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) & (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) & (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) & (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) & (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))) | (((e23 = h9(e13)) | (e23 = h9(e12)) | (e23 = h9(e11)) | (e23 = h9(e10))) & ((e22 = h9(e13)) | (e22 = h9(e12)) | (e22 = h9(e11)) | (e22 = h9(e10))) & ((e21 = h9(e13)) | (e21 = h9(e12)) | (e21 = h9(e11)) | (e21 = h9(e10))) & ((e20 = h9(e13)) | (e20 = h9(e12)) | (e20 = h9(e11)) | (e20 = h9(e10))) & (h9(op1(e13, e13)) = op2(h9(e13), h9(e13))) & (h9(op1(e13, e12)) = op2(h9(e13), h9(e12))) & (h9(op1(e13, e11)) = op2(h9(e13), h9(e11))) & (h9(op1(e13, e10)) = op2(h9(e13), h9(e10))) & (h9(op1(e12, e13)) = op2(h9(e12), h9(e13))) & (h9(op1(e12, e12)) = op2(h9(e12), h9(e12))) & (h9(op1(e12, e11)) = op2(h9(e12), h9(e11))) & (h9(op1(e12, e10)) = op2(h9(e12), h9(e10))) & (h9(op1(e11, e13)) = op2(h9(e11), h9(e13))) & (h9(op1(e11, e12)) = op2(h9(e11), h9(e12))) & (h9(op1(e11, e11)) = op2(h9(e11), h9(e11))) & (h9(op1(e11, e10)) = op2(h9(e11), h9(e10))) & (h9(op1(e10, e13)) = op2(h9(e10), h9(e13))) & (h9(op1(e10, e12)) = op2(h9(e10), h9(e12))) & (h9(op1(e10, e11)) = op2(h9(e10), h9(e11))) & (h9(op1(e10, e10)) = op2(h9(e10), h9(e10)))) | (((e23 = h8(e13)) | (e23 = h8(e12)) | (e23 = h8(e11)) | (e23 = h8(e10))) & ((e22 = h8(e13)) | (e22 = h8(e12)) | (e22 = h8(e11)) | (e22 = h8(e10))) & ((e21 = h8(e13)) | (e21 = h8(e12)) | (e21 = h8(e11)) | (e21 = h8(e10))) & ((e20 = h8(e13)) | (e20 = h8(e12)) | (e20 = h8(e11)) | (e20 = h8(e10))) & (h8(op1(e13, e13)) = op2(h8(e13), h8(e13))) & (h8(op1(e13, e12)) = op2(h8(e13), h8(e12))) & (h8(op1(e13, e11)) = op2(h8(e13), h8(e11))) & (h8(op1(e13, e10)) = op2(h8(e13), h8(e10))) & (h8(op1(e12, e13)) = op2(h8(e12), h8(e13))) & (h8(op1(e12, e12)) = op2(h8(e12), h8(e12))) & (h8(op1(e12, e11)) = op2(h8(e12), h8(e11))) & (h8(op1(e12, e10)) = op2(h8(e12), h8(e10))) & (h8(op1(e11, e13)) = op2(h8(e11), h8(e13))) & (h8(op1(e11, e12)) = op2(h8(e11), h8(e12))) & (h8(op1(e11, e11)) = op2(h8(e11), h8(e11))) & (h8(op1(e11, e10)) = op2(h8(e11), h8(e10))) & (h8(op1(e10, e13)) = op2(h8(e10), h8(e13))) & (h8(op1(e10, e12)) = op2(h8(e10), h8(e12))) & (h8(op1(e10, e11)) = op2(h8(e10), h8(e11))) & (h8(op1(e10, e10)) = op2(h8(e10), h8(e10)))) | (((e23 = h7(e13)) | (e23 = h7(e12)) | (e23 = h7(e11)) | (e23 = h7(e10))) & ((e22 = h7(e13)) | (e22 = h7(e12)) | (e22 = h7(e11)) | (e22 = h7(e10))) & ((e21 = h7(e13)) | (e21 = h7(e12)) | (e21 = h7(e11)) | (e21 = h7(e10))) & ((e20 = h7(e13)) | (e20 = h7(e12)) | (e20 = h7(e11)) | (e20 = h7(e10))) & (h7(op1(e13, e13)) = op2(h7(e13), h7(e13))) & (h7(op1(e13, e12)) = op2(h7(e13), h7(e12))) & (h7(op1(e13, e11)) = op2(h7(e13), h7(e11))) & (h7(op1(e13, e10)) = op2(h7(e13), h7(e10))) & (h7(op1(e12, e13)) = op2(h7(e12), h7(e13))) & (h7(op1(e12, e12)) = op2(h7(e12), h7(e12))) & (h7(op1(e12, e11)) = op2(h7(e12), h7(e11))) & (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) & (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) & (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) & (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) & (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) & (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) & (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) & (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) & (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))) | (((e23 = h6(e13)) | (e23 = h6(e12)) | (e23 = h6(e11)) | (e23 = h6(e10))) & ((e22 = h6(e13)) | (e22 = h6(e12)) | (e22 = h6(e11)) | (e22 = h6(e10))) & ((e21 = h6(e13)) | (e21 = h6(e12)) | (e21 = h6(e11)) | (e21 = h6(e10))) & ((e20 = h6(e13)) | (e20 = h6(e12)) | (e20 = h6(e11)) | (e20 = h6(e10))) & (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) & (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) & (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) & (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) & (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) & (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) & (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) & (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) & (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) & (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) & (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) & (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) & (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) & (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) & (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) & (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))) | (((e23 = h5(e13)) | (e23 = h5(e12)) | (e23 = h5(e11)) | (e23 = h5(e10))) & ((e22 = h5(e13)) | (e22 = h5(e12)) | (e22 = h5(e11)) | (e22 = h5(e10))) & ((e21 = h5(e13)) | (e21 = h5(e12)) | (e21 = h5(e11)) | (e21 = h5(e10))) & ((e20 = h5(e13)) | (e20 = h5(e12)) | (e20 = h5(e11)) | (e20 = h5(e10))) & (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) & (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) & (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) & (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) & (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) & (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) & (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) & (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) & (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) & (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) & (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) & (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) & (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) & (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) & (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) & (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))) | (((e23 = h4(e13)) | (e23 = h4(e12)) | (e23 = h4(e11)) | (e23 = h4(e10))) & ((e22 = h4(e13)) | (e22 = h4(e12)) | (e22 = h4(e11)) | (e22 = h4(e10))) & ((e21 = h4(e13)) | (e21 = h4(e12)) | (e21 = h4(e11)) | (e21 = h4(e10))) & ((e20 = h4(e13)) | (e20 = h4(e12)) | (e20 = h4(e11)) | (e20 = h4(e10))) & (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) & (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) & (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) & (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) & (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) & (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) & (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) & (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) & (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) & (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) & (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) & (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) & (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) & (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) & (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) & (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) | (((e23 = h3(e13)) | (e23 = h3(e12)) | (e23 = h3(e11)) | (e23 = h3(e10))) & ((e22 = h3(e13)) | (e22 = h3(e12)) | (e22 = h3(e11)) | (e22 = h3(e10))) & ((e21 = h3(e13)) | (e21 = h3(e12)) | (e21 = h3(e11)) | (e21 = h3(e10))) & ((e20 = h3(e13)) | (e20 = h3(e12)) | (e20 = h3(e11)) | (e20 = h3(e10))) & (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) & (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) & (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) & (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) & (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) & (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) & (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) & (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) & (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) & (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) & (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) & (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) & (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) & (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) & (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) & (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) | (((e23 = h2(e13)) | (e23 = h2(e12)) | (e23 = h2(e11)) | (e23 = h2(e10))) & ((e22 = h2(e13)) | (e22 = h2(e12)) | (e22 = h2(e11)) | (e22 = h2(e10))) & ((e21 = h2(e13)) | (e21 = h2(e12)) | (e21 = h2(e11)) | (e21 = h2(e10))) & ((e20 = h2(e13)) | (e20 = h2(e12)) | (e20 = h2(e11)) | (e20 = h2(e10))) & (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) & (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) & (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) & (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) & (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) & (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) & (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) & (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) & (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) & (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) & (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) & (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) & (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) & (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) & (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) & (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) | (((e23 = h1(e13)) | (e23 = h1(e12)) | (e23 = h1(e11)) | (e23 = h1(e10))) & ((e22 = h1(e13)) | (e22 = h1(e12)) | (e22 = h1(e11)) | (e22 = h1(e10))) & ((e21 = h1(e13)) | (e21 = h1(e12)) | (e21 = h1(e11)) | (e21 = h1(e10))) & ((e20 = h1(e13)) | (e20 = h1(e12)) | (e20 = h1(e11)) | (e20 = h1(e10))) & (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) & (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) & (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) & (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) & (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) & (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) & (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) & (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) & (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) & (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) & (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) & (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) & (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) & (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) & (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) & (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG037+1.p', co1)).
fof(f1459, plain, (~ spl42_152 | ~ spl42_148), inference(avatar_split_clause, [], [f1458, f1425, f1449])).
fof(f1458, plain, (~ (e20 = h10(e10)) | ~ sP36), inference(forward_demodulation, [], [f460, f1392])).
fof(f460, plain, (~ (e20 = h11(e10)) | ~ sP36), inference(cnf_transformation, [], [f87])).
fof(f87, plain, ((~ (e20 = h11(e13)) & ~ (e20 = h11(e12)) & ~ (e20 = h11(e11)) & ~ (e20 = h11(e10))) | ~ sP36), inference(nnf_transformation, [], [f69])).
fof(f1438, plain, (~ spl42_149 | ~ spl42_150), inference(avatar_split_clause, [], [f454, f1435, f1431])).
fof(f454, plain, (~ (e22 = h11(e12)) | ~ sP38), inference(cnf_transformation, [], [f85])).
fof(f85, plain, ((~ (e22 = h11(e13)) & ~ (e22 = h11(e12)) & ~ (e22 = h11(e11)) & ~ (e22 = h11(e10))) | ~ sP38), inference(nnf_transformation, [], [f71])).
fof(f1333, plain, spl42_109, inference(avatar_split_clause, [], [f390, f1124])).
fof(f390, plain, (e20 = op2(e21, e21)), inference(cnf_transformation, [], [f15])).
fof(f15, plain, ((e22 = op2(e23, e21)) & (e20 = op2(e21, e21))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG037+1.p', ax15)).
fof(f1332, plain, spl42_79, inference(avatar_split_clause, [], [f391, f996])).
fof(f391, plain, (e22 = op2(e23, e21)), inference(cnf_transformation, [], [f15])).
fof(f1331, plain, spl42_41, inference(avatar_split_clause, [], [f388, f803])).
fof(f388, plain, (e10 = op1(e11, e11)), inference(cnf_transformation, [], [f14])).
fof(f14, plain, ((e12 = op1(e13, e11)) & (e10 = op1(e11, e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG037+1.p', ax14)).
fof(f1330, plain, spl42_11, inference(avatar_split_clause, [], [f389, f675])).
fof(f389, plain, (e12 = op1(e13, e11)), inference(cnf_transformation, [], [f14])).
fof(f1328, plain, (spl42_142 | spl42_141 | spl42_140 | spl42_112), inference(avatar_split_clause, [], [f385, f1136, f1303, f1311, f1319])).
fof(f1319, plain, (spl42_142 <=> sP3), introduced(avatar_definition, [new_symbols(naming, [spl42_142])])).
fof(f1311, plain, (spl42_141 <=> sP4), introduced(avatar_definition, [new_symbols(naming, [spl42_141])])).
fof(f1303, plain, (spl42_140 <=> sP5), introduced(avatar_definition, [new_symbols(naming, [spl42_140])])).
fof(f385, plain, ((e23 = op2(e21, e21)) | sP5 | sP4 | sP3), inference(cnf_transformation, [], [f38])).
fof(f38, plain, (((e23 = op2(e23, e23)) & (e23 = op2(e22, e22)) & (e23 = op2(e21, e21)) & (op2(e20, e20) = e23)) | sP5 | sP4 | sP3), inference(definition_folding, [], [f13, e37, e36, e35])).
fof(f35, plain, (((e20 = op2(e23, e23)) & (e20 = op2(e22, e22)) & (e20 = op2(e21, e21)) & (e20 = op2(e20, e20))) | ~ sP3), inference(usedef, [], [e35])).
fof(e35, plain, (sP3 <=> ((e20 = op2(e23, e23)) & (e20 = op2(e22, e22)) & (e20 = op2(e21, e21)) & (e20 = op2(e20, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f36, plain, (((e21 = op2(e23, e23)) & (e21 = op2(e22, e22)) & (e21 = op2(e21, e21)) & (op2(e20, e20) = e21)) | ~ sP4), inference(usedef, [], [e36])).
fof(e36, plain, (sP4 <=> ((e21 = op2(e23, e23)) & (e21 = op2(e22, e22)) & (e21 = op2(e21, e21)) & (op2(e20, e20) = e21))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f37, plain, (((e22 = op2(e23, e23)) & (e22 = op2(e22, e22)) & (e22 = op2(e21, e21)) & (op2(e20, e20) = e22)) | ~ sP5), inference(usedef, [], [e37])).
fof(e37, plain, (sP5 <=> ((e22 = op2(e23, e23)) & (e22 = op2(e22, e22)) & (e22 = op2(e21, e21)) & (op2(e20, e20) = e22))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f13, plain, (((e23 = op2(e23, e23)) & (e23 = op2(e22, e22)) & (e23 = op2(e21, e21)) & (op2(e20, e20) = e23)) | ((e22 = op2(e23, e23)) & (e22 = op2(e22, e22)) & (e22 = op2(e21, e21)) & (op2(e20, e20) = e22)) | ((e21 = op2(e23, e23)) & (e21 = op2(e22, e22)) & (e21 = op2(e21, e21)) & (op2(e20, e20) = e21)) | ((e20 = op2(e23, e23)) & (e20 = op2(e22, e22)) & (e20 = op2(e21, e21)) & (e20 = op2(e20, e20)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG037+1.p', ax13)).
fof(f1325, plain, (~ spl42_142 | spl42_129), inference(avatar_split_clause, [], [f380, f1209, f1319])).
fof(f380, plain, ((e20 = op2(e20, e20)) | ~ sP3), inference(cnf_transformation, [], [f81])).
fof(f81, plain, (((e20 = op2(e23, e23)) & (e20 = op2(e22, e22)) & (e20 = op2(e21, e21)) & (e20 = op2(e20, e20))) | ~ sP3), inference(nnf_transformation, [], [f35])).
fof(f1323, plain, (~ spl42_142 | spl42_89), inference(avatar_split_clause, [], [f382, f1039, f1319])).
fof(f382, plain, ((e20 = op2(e22, e22)) | ~ sP3), inference(cnf_transformation, [], [f81])).
fof(f1322, plain, (~ spl42_142 | spl42_69), inference(avatar_split_clause, [], [f383, f954, f1319])).
fof(f383, plain, ((e20 = op2(e23, e23)) | ~ sP3), inference(cnf_transformation, [], [f81])).
fof(f1316, plain, (~ spl42_141 | spl42_110), inference(avatar_split_clause, [], [f377, f1128, f1311])).
fof(f377, plain, ((e21 = op2(e21, e21)) | ~ sP4), inference(cnf_transformation, [], [f80])).
fof(f80, plain, (((e21 = op2(e23, e23)) & (e21 = op2(e22, e22)) & (e21 = op2(e21, e21)) & (op2(e20, e20) = e21)) | ~ sP4), inference(nnf_transformation, [], [f36])).
fof(f1308, plain, (~ spl42_140 | spl42_111), inference(avatar_split_clause, [], [f373, f1132, f1303])).
fof(f373, plain, ((e22 = op2(e21, e21)) | ~ sP5), inference(cnf_transformation, [], [f79])).
fof(f79, plain, (((e22 = op2(e23, e23)) & (e22 = op2(e22, e22)) & (e22 = op2(e21, e21)) & (op2(e20, e20) = e22)) | ~ sP5), inference(nnf_transformation, [], [f37])).
fof(f1300, plain, (spl42_139 | spl42_138 | spl42_137 | spl42_44), inference(avatar_split_clause, [], [f369, f815, f1275, f1283, f1291])).
fof(f1291, plain, (spl42_139 <=> sP0), introduced(avatar_definition, [new_symbols(naming, [spl42_139])])).
fof(f1283, plain, (spl42_138 <=> sP1), introduced(avatar_definition, [new_symbols(naming, [spl42_138])])).
fof(f1275, plain, (spl42_137 <=> sP2), introduced(avatar_definition, [new_symbols(naming, [spl42_137])])).
fof(f369, plain, ((e13 = op1(e11, e11)) | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f34])).
fof(f34, plain, (((e13 = op1(e13, e13)) & (e13 = op1(e12, e12)) & (e13 = op1(e11, e11)) & (op1(e10, e10) = e13)) | sP2 | sP1 | sP0), inference(definition_folding, [], [f12, e33, e32, e31])).
fof(f31, plain, (((e10 = op1(e13, e13)) & (e10 = op1(e12, e12)) & (e10 = op1(e11, e11)) & (e10 = op1(e10, e10))) | ~ sP0), inference(usedef, [], [e31])).
fof(e31, plain, (sP0 <=> ((e10 = op1(e13, e13)) & (e10 = op1(e12, e12)) & (e10 = op1(e11, e11)) & (e10 = op1(e10, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f32, plain, (((e11 = op1(e13, e13)) & (e11 = op1(e12, e12)) & (e11 = op1(e11, e11)) & (op1(e10, e10) = e11)) | ~ sP1), inference(usedef, [], [e32])).
fof(e32, plain, (sP1 <=> ((e11 = op1(e13, e13)) & (e11 = op1(e12, e12)) & (e11 = op1(e11, e11)) & (op1(e10, e10) = e11))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f33, plain, (((e12 = op1(e13, e13)) & (e12 = op1(e12, e12)) & (e12 = op1(e11, e11)) & (op1(e10, e10) = e12)) | ~ sP2), inference(usedef, [], [e33])).
fof(e33, plain, (sP2 <=> ((e12 = op1(e13, e13)) & (e12 = op1(e12, e12)) & (e12 = op1(e11, e11)) & (op1(e10, e10) = e12))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f12, plain, (((e13 = op1(e13, e13)) & (e13 = op1(e12, e12)) & (e13 = op1(e11, e11)) & (op1(e10, e10) = e13)) | ((e12 = op1(e13, e13)) & (e12 = op1(e12, e12)) & (e12 = op1(e11, e11)) & (op1(e10, e10) = e12)) | ((e11 = op1(e13, e13)) & (e11 = op1(e12, e12)) & (e11 = op1(e11, e11)) & (op1(e10, e10) = e11)) | ((e10 = op1(e13, e13)) & (e10 = op1(e12, e12)) & (e10 = op1(e11, e11)) & (e10 = op1(e10, e10)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG037+1.p', ax12)).
fof(f1297, plain, (~ spl42_139 | spl42_61), inference(avatar_split_clause, [], [f364, f888, f1291])).
fof(f364, plain, ((e10 = op1(e10, e10)) | ~ sP0), inference(cnf_transformation, [], [f78])).
fof(f78, plain, (((e10 = op1(e13, e13)) & (e10 = op1(e12, e12)) & (e10 = op1(e11, e11)) & (e10 = op1(e10, e10))) | ~ sP0), inference(nnf_transformation, [], [f31])).
fof(f1295, plain, (~ spl42_139 | spl42_21), inference(avatar_split_clause, [], [f366, f718, f1291])).
fof(f366, plain, ((e10 = op1(e12, e12)) | ~ sP0), inference(cnf_transformation, [], [f78])).
fof(f1294, plain, (~ spl42_139 | spl42_1), inference(avatar_split_clause, [], [f367, f633, f1291])).
fof(f367, plain, ((e10 = op1(e13, e13)) | ~ sP0), inference(cnf_transformation, [], [f78])).
fof(f1288, plain, (~ spl42_138 | spl42_42), inference(avatar_split_clause, [], [f361, f807, f1283])).
fof(f361, plain, ((e11 = op1(e11, e11)) | ~ sP1), inference(cnf_transformation, [], [f77])).
fof(f77, plain, (((e11 = op1(e13, e13)) & (e11 = op1(e12, e12)) & (e11 = op1(e11, e11)) & (op1(e10, e10) = e11)) | ~ sP1), inference(nnf_transformation, [], [f32])).
fof(f1278, plain, (~ spl42_137 | spl42_3), inference(avatar_split_clause, [], [f359, f641, f1275])).
fof(f359, plain, ((e12 = op1(e13, e13)) | ~ sP2), inference(cnf_transformation, [], [f76])).
fof(f76, plain, (((e12 = op1(e13, e13)) & (e12 = op1(e12, e12)) & (e12 = op1(e11, e11)) & (op1(e10, e10) = e12)) | ~ sP2), inference(nnf_transformation, [], [f33])).
fof(f1259, plain, (spl42_116 | spl42_112 | spl42_108 | spl42_104), inference(avatar_split_clause, [], [f214, f1102, f1119, f1136, f1153])).
fof(f214, plain, ((e23 = op2(e21, e23)) | (e23 = op2(e21, e22)) | (e23 = op2(e21, e21)) | (e23 = op2(e21, e20))), inference(cnf_transformation, [], [f6])).
fof(f6, plain, (((e23 = op2(e23, e23)) | (e23 = op2(e22, e23)) | (e23 = op2(e21, e23)) | (e23 = op2(e20, e23))) & ((e23 = op2(e23, e23)) | (e23 = op2(e23, e22)) | (e23 = op2(e23, e21)) | (e23 = op2(e23, e20))) & ((e22 = op2(e23, e23)) | (e22 = op2(e22, e23)) | (e22 = op2(e21, e23)) | (e22 = op2(e20, e23))) & ((e22 = op2(e23, e23)) | (e22 = op2(e23, e22)) | (e22 = op2(e23, e21)) | (e22 = op2(e23, e20))) & ((e21 = op2(e23, e23)) | (e21 = op2(e22, e23)) | (e21 = op2(e21, e23)) | (e21 = op2(e20, e23))) & ((e21 = op2(e23, e23)) | (e21 = op2(e23, e22)) | (e21 = op2(e23, e21)) | (e21 = op2(e23, e20))) & ((e20 = op2(e23, e23)) | (e20 = op2(e22, e23)) | (e20 = op2(e21, e23)) | (e20 = op2(e20, e23))) & ((e20 = op2(e23, e23)) | (e20 = op2(e23, e22)) | (e20 = op2(e23, e21)) | (e20 = op2(e23, e20))) & ((e23 = op2(e23, e22)) | (e23 = op2(e22, e22)) | (e23 = op2(e21, e22)) | (e23 = op2(e20, e22))) & ((e23 = op2(e22, e23)) | (e23 = op2(e22, e22)) | (e23 = op2(e22, e21)) | (e23 = op2(e22, e20))) & ((e22 = op2(e23, e22)) | (e22 = op2(e22, e22)) | (e22 = op2(e21, e22)) | (e22 = op2(e20, e22))) & ((e22 = op2(e22, e23)) | (e22 = op2(e22, e22)) | (e22 = op2(e22, e21)) | (e22 = op2(e22, e20))) & ((e21 = op2(e23, e22)) | (e21 = op2(e22, e22)) | (e21 = op2(e21, e22)) | (e21 = op2(e20, e22))) & ((e21 = op2(e22, e23)) | (e21 = op2(e22, e22)) | (e21 = op2(e22, e21)) | (e21 = op2(e22, e20))) & ((e20 = op2(e23, e22)) | (e20 = op2(e22, e22)) | (e20 = op2(e21, e22)) | (e20 = op2(e20, e22))) & ((e20 = op2(e22, e23)) | (e20 = op2(e22, e22)) | (e20 = op2(e22, e21)) | (e20 = op2(e22, e20))) & ((e23 = op2(e23, e21)) | (e23 = op2(e22, e21)) | (e23 = op2(e21, e21)) | (e23 = op2(e20, e21))) & ((e23 = op2(e21, e23)) | (e23 = op2(e21, e22)) | (e23 = op2(e21, e21)) | (e23 = op2(e21, e20))) & ((e22 = op2(e23, e21)) | (e22 = op2(e22, e21)) | (e22 = op2(e21, e21)) | (e22 = op2(e20, e21))) & ((e22 = op2(e21, e23)) | (e22 = op2(e21, e22)) | (e22 = op2(e21, e21)) | (e22 = op2(e21, e20))) & ((e21 = op2(e23, e21)) | (e21 = op2(e22, e21)) | (e21 = op2(e21, e21)) | (e21 = op2(e20, e21))) & ((e21 = op2(e21, e23)) | (e21 = op2(e21, e22)) | (e21 = op2(e21, e21)) | (e21 = op2(e21, e20))) & ((e20 = op2(e23, e21)) | (e20 = op2(e22, e21)) | (e20 = op2(e21, e21)) | (e20 = op2(e20, e21))) & ((e20 = op2(e21, e23)) | (e20 = op2(e21, e22)) | (e20 = op2(e21, e21)) | (e20 = op2(e21, e20))) & ((e23 = op2(e23, e20)) | (e23 = op2(e22, e20)) | (e23 = op2(e21, e20)) | (op2(e20, e20) = e23)) & ((e23 = op2(e20, e23)) | (e23 = op2(e20, e22)) | (e23 = op2(e20, e21)) | (op2(e20, e20) = e23)) & ((e22 = op2(e23, e20)) | (e22 = op2(e22, e20)) | (e22 = op2(e21, e20)) | (op2(e20, e20) = e22)) & ((e22 = op2(e20, e23)) | (e22 = op2(e20, e22)) | (e22 = op2(e20, e21)) | (op2(e20, e20) = e22)) & ((e21 = op2(e23, e20)) | (e21 = op2(e22, e20)) | (e21 = op2(e21, e20)) | (op2(e20, e20) = e21)) & ((e21 = op2(e20, e23)) | (e21 = op2(e20, e22)) | (e21 = op2(e20, e21)) | (op2(e20, e20) = e21)) & ((e20 = op2(e23, e20)) | (e20 = op2(e22, e20)) | (e20 = op2(e21, e20)) | (e20 = op2(e20, e20))) & ((e20 = op2(e20, e23)) | (e20 = op2(e20, e22)) | (e20 = op2(e20, e21)) | (e20 = op2(e20, e20)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG037+1.p', ax6)).
fof(f1258, plain, (spl42_128 | spl42_112 | spl42_96 | spl42_80), inference(avatar_split_clause, [], [f215, f1000, f1068, f1136, f1204])).
fof(f215, plain, ((e23 = op2(e23, e21)) | (e23 = op2(e22, e21)) | (e23 = op2(e21, e21)) | (e23 = op2(e20, e21))), inference(cnf_transformation, [], [f6])).
fof(f1247, plain, (spl42_82 | spl42_78 | spl42_74 | spl42_70), inference(avatar_split_clause, [], [f226, f958, f975, f992, f1009])).
fof(f226, plain, ((e21 = op2(e23, e23)) | (e21 = op2(e23, e22)) | (e21 = op2(e23, e21)) | (e21 = op2(e23, e20))), inference(cnf_transformation, [], [f6])).
fof(f1246, plain, (spl42_118 | spl42_102 | spl42_86 | spl42_70), inference(avatar_split_clause, [], [f227, f958, f1026, f1094, f1162])).
fof(f227, plain, ((e21 = op2(e23, e23)) | (e21 = op2(e22, e23)) | (e21 = op2(e21, e23)) | (e21 = op2(e20, e23))), inference(cnf_transformation, [], [f6])).
fof(f1244, plain, (spl42_119 | spl42_103 | spl42_87 | spl42_71), inference(avatar_split_clause, [], [f229, f962, f1030, f1098, f1166])).
fof(f229, plain, ((e22 = op2(e23, e23)) | (e22 = op2(e22, e23)) | (e22 = op2(e21, e23)) | (e22 = op2(e20, e23))), inference(cnf_transformation, [], [f6])).
fof(f1241, plain, (spl42_133 | spl42_134 | spl42_135 | spl42_136), inference(avatar_split_clause, [], [f199, f1238, f1234, f1230, f1226])).
fof(f199, plain, ((e23 = unit2) | (e22 = unit2) | (e21 = unit2) | (e20 = unit2)), inference(cnf_transformation, [], [f5])).
fof(f937, plain, (spl42_60 | spl42_44 | spl42_28 | spl42_12), inference(avatar_split_clause, [], [f158, f679, f747, f815, f883])).
fof(f158, plain, ((e13 = op1(e13, e11)) | (e13 = op1(e12, e11)) | (e13 = op1(e11, e11)) | (e13 = op1(e10, e11))), inference(cnf_transformation, [], [f3])).
fof(f3, plain, (((e13 = op1(e13, e13)) | (e13 = op1(e12, e13)) | (e13 = op1(e11, e13)) | (e13 = op1(e10, e13))) & ((e13 = op1(e13, e13)) | (e13 = op1(e13, e12)) | (e13 = op1(e13, e11)) | (e13 = op1(e13, e10))) & ((e12 = op1(e13, e13)) | (e12 = op1(e12, e13)) | (e12 = op1(e11, e13)) | (e12 = op1(e10, e13))) & ((e12 = op1(e13, e13)) | (e12 = op1(e13, e12)) | (e12 = op1(e13, e11)) | (e12 = op1(e13, e10))) & ((e11 = op1(e13, e13)) | (e11 = op1(e12, e13)) | (e11 = op1(e11, e13)) | (e11 = op1(e10, e13))) & ((e11 = op1(e13, e13)) | (e11 = op1(e13, e12)) | (e11 = op1(e13, e11)) | (e11 = op1(e13, e10))) & ((e10 = op1(e13, e13)) | (e10 = op1(e12, e13)) | (e10 = op1(e11, e13)) | (e10 = op1(e10, e13))) & ((e10 = op1(e13, e13)) | (e10 = op1(e13, e12)) | (e10 = op1(e13, e11)) | (e10 = op1(e13, e10))) & ((e13 = op1(e13, e12)) | (e13 = op1(e12, e12)) | (e13 = op1(e11, e12)) | (e13 = op1(e10, e12))) & ((e13 = op1(e12, e13)) | (e13 = op1(e12, e12)) | (e13 = op1(e12, e11)) | (e13 = op1(e12, e10))) & ((e12 = op1(e13, e12)) | (e12 = op1(e12, e12)) | (e12 = op1(e11, e12)) | (e12 = op1(e10, e12))) & ((e12 = op1(e12, e13)) | (e12 = op1(e12, e12)) | (e12 = op1(e12, e11)) | (e12 = op1(e12, e10))) & ((e11 = op1(e13, e12)) | (e11 = op1(e12, e12)) | (e11 = op1(e11, e12)) | (e11 = op1(e10, e12))) & ((e11 = op1(e12, e13)) | (e11 = op1(e12, e12)) | (e11 = op1(e12, e11)) | (e11 = op1(e12, e10))) & ((e10 = op1(e13, e12)) | (e10 = op1(e12, e12)) | (e10 = op1(e11, e12)) | (e10 = op1(e10, e12))) & ((e10 = op1(e12, e13)) | (e10 = op1(e12, e12)) | (e10 = op1(e12, e11)) | (e10 = op1(e12, e10))) & ((e13 = op1(e13, e11)) | (e13 = op1(e12, e11)) | (e13 = op1(e11, e11)) | (e13 = op1(e10, e11))) & ((e13 = op1(e11, e13)) | (e13 = op1(e11, e12)) | (e13 = op1(e11, e11)) | (e13 = op1(e11, e10))) & ((e12 = op1(e13, e11)) | (e12 = op1(e12, e11)) | (e12 = op1(e11, e11)) | (e12 = op1(e10, e11))) & ((e12 = op1(e11, e13)) | (e12 = op1(e11, e12)) | (e12 = op1(e11, e11)) | (e12 = op1(e11, e10))) & ((e11 = op1(e13, e11)) | (e11 = op1(e12, e11)) | (e11 = op1(e11, e11)) | (e11 = op1(e10, e11))) & ((e11 = op1(e11, e13)) | (e11 = op1(e11, e12)) | (e11 = op1(e11, e11)) | (e11 = op1(e11, e10))) & ((e10 = op1(e13, e11)) | (e10 = op1(e12, e11)) | (e10 = op1(e11, e11)) | (e10 = op1(e10, e11))) & ((e10 = op1(e11, e13)) | (e10 = op1(e11, e12)) | (e10 = op1(e11, e11)) | (e10 = op1(e11, e10))) & ((e13 = op1(e13, e10)) | (e13 = op1(e12, e10)) | (e13 = op1(e11, e10)) | (op1(e10, e10) = e13)) & ((e13 = op1(e10, e13)) | (e13 = op1(e10, e12)) | (e13 = op1(e10, e11)) | (op1(e10, e10) = e13)) & ((e12 = op1(e13, e10)) | (e12 = op1(e12, e10)) | (e12 = op1(e11, e10)) | (op1(e10, e10) = e12)) & ((e12 = op1(e10, e13)) | (e12 = op1(e10, e12)) | (e12 = op1(e10, e11)) | (op1(e10, e10) = e12)) & ((e11 = op1(e13, e10)) | (e11 = op1(e12, e10)) | (e11 = op1(e11, e10)) | (op1(e10, e10) = e11)) & ((e11 = op1(e10, e13)) | (e11 = op1(e10, e12)) | (e11 = op1(e10, e11)) | (op1(e10, e10) = e11)) & ((e10 = op1(e13, e10)) | (e10 = op1(e12, e10)) | (e10 = op1(e11, e10)) | (e10 = op1(e10, e10))) & ((e10 = op1(e10, e13)) | (e10 = op1(e10, e12)) | (e10 = op1(e10, e11)) | (e10 = op1(e10, e10)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG037+1.p', ax3)).
fof(f926, plain, (spl42_14 | spl42_10 | spl42_6 | spl42_2), inference(avatar_split_clause, [], [f169, f637, f654, f671, f688])).
fof(f169, plain, ((e11 = op1(e13, e13)) | (e11 = op1(e13, e12)) | (e11 = op1(e13, e11)) | (e11 = op1(e13, e10))), inference(cnf_transformation, [], [f3])).
fof(f925, plain, (spl42_50 | spl42_34 | spl42_18 | spl42_2), inference(avatar_split_clause, [], [f170, f637, f705, f773, f841])).
fof(f170, plain, ((e11 = op1(e13, e13)) | (e11 = op1(e12, e13)) | (e11 = op1(e11, e13)) | (e11 = op1(e10, e13))), inference(cnf_transformation, [], [f3])).
fof(f923, plain, (spl42_51 | spl42_35 | spl42_19 | spl42_3), inference(avatar_split_clause, [], [f172, f641, f709, f777, f845])).
fof(f172, plain, ((e12 = op1(e13, e13)) | (e12 = op1(e12, e13)) | (e12 = op1(e11, e13)) | (e12 = op1(e10, e13))), inference(cnf_transformation, [], [f3])).
fof(f920, plain, (spl42_65 | spl42_66 | spl42_67 | spl42_68), inference(avatar_split_clause, [], [f142, f917, f913, f909, f905])).
fof(f142, plain, ((e13 = unit1) | (e12 = unit1) | (e11 = unit1) | (e10 = unit1)), inference(cnf_transformation, [], [f2])).
fof(f801, plain, (spl42_37 | spl42_38 | spl42_39 | spl42_40), inference(avatar_split_clause, [], [f124, f798, f794, f790, f786])).
fof(f124, plain, ((e13 = op1(e11, e12)) | (e12 = op1(e11, e12)) | (e11 = op1(e11, e12)) | (e10 = op1(e11, e12))), inference(cnf_transformation, [], [f1])).
fof(f1, plain, (((e13 = op1(e13, e13)) | (e12 = op1(e13, e13)) | (e11 = op1(e13, e13)) | (e10 = op1(e13, e13))) & ((e13 = op1(e13, e12)) | (e12 = op1(e13, e12)) | (e11 = op1(e13, e12)) | (e10 = op1(e13, e12))) & ((e13 = op1(e13, e11)) | (e12 = op1(e13, e11)) | (e11 = op1(e13, e11)) | (e10 = op1(e13, e11))) & ((e13 = op1(e13, e10)) | (e12 = op1(e13, e10)) | (e11 = op1(e13, e10)) | (e10 = op1(e13, e10))) & ((e13 = op1(e12, e13)) | (e12 = op1(e12, e13)) | (e11 = op1(e12, e13)) | (e10 = op1(e12, e13))) & ((e13 = op1(e12, e12)) | (e12 = op1(e12, e12)) | (e11 = op1(e12, e12)) | (e10 = op1(e12, e12))) & ((e13 = op1(e12, e11)) | (e12 = op1(e12, e11)) | (e11 = op1(e12, e11)) | (e10 = op1(e12, e11))) & ((e13 = op1(e12, e10)) | (e12 = op1(e12, e10)) | (e11 = op1(e12, e10)) | (e10 = op1(e12, e10))) & ((e13 = op1(e11, e13)) | (e12 = op1(e11, e13)) | (e11 = op1(e11, e13)) | (e10 = op1(e11, e13))) & ((e13 = op1(e11, e12)) | (e12 = op1(e11, e12)) | (e11 = op1(e11, e12)) | (e10 = op1(e11, e12))) & ((e13 = op1(e11, e11)) | (e12 = op1(e11, e11)) | (e11 = op1(e11, e11)) | (e10 = op1(e11, e11))) & ((e13 = op1(e11, e10)) | (e12 = op1(e11, e10)) | (e11 = op1(e11, e10)) | (e10 = op1(e11, e10))) & ((e13 = op1(e10, e13)) | (e12 = op1(e10, e13)) | (e11 = op1(e10, e13)) | (e10 = op1(e10, e13))) & ((e13 = op1(e10, e12)) | (e12 = op1(e10, e12)) | (e11 = op1(e10, e12)) | (e10 = op1(e10, e12))) & ((e13 = op1(e10, e11)) | (e12 = op1(e10, e11)) | (e11 = op1(e10, e11)) | (e10 = op1(e10, e11))) & ((op1(e10, e10) = e13) | (op1(e10, e10) = e12) | (op1(e10, e10) = e11) | (e10 = op1(e10, e10)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG037+1.p', ax1)).