fof(f1503, plain, $false, inference(subsumption_resolution, [], [f1502, f1075])).
fof(f1075, plain, ~ sP0, inference(subsumption_resolution, [], [f1074, f168])).
fof(f168, plain, ~ (e2 = e4), inference(cnf_transformation, [], [f1])).
fof(f1, plain, (~ (e3 = e4) & ~ (e2 = e4) & ~ (e2 = e3) & ~ (e1 = e4) & ~ (e1 = e3) & ~ (e1 = e2) & ~ (e0 = e4) & ~ (e0 = e3) & ~ (e0 = e2) & ~ (e0 = e1)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG094+1.p', ax1)).
fof(f1074, plain, ((e2 = e4) | ~ sP0), inference(forward_demodulation, [], [f1073, f181])).
fof(f181, plain, (e4 = op(e2, e1)), inference(cnf_transformation, [], [f2])).
fof(f2, plain, ((e1 = op(e4, e4)) & (e2 = op(e4, e3)) & (e0 = op(e4, e2)) & (e3 = op(e4, e1)) & (e4 = op(e4, e0)) & (e2 = op(e3, e4)) & (e1 = op(e3, e3)) & (e4 = op(e3, e2)) & (e0 = op(e3, e1)) & (e3 = op(e3, e0)) & (e3 = op(e2, e4)) & (e0 = op(e2, e3)) & (e1 = op(e2, e2)) & (e4 = op(e2, e1)) & (e2 = op(e2, e0)) & (e0 = op(e1, e4)) & (e4 = op(e1, e3)) & (e3 = op(e1, e2)) & (e2 = op(e1, e1)) & (e1 = op(e1, e0)) & (e4 = op(e0, e4)) & (e3 = op(e0, e3)) & (e2 = op(e0, e2)) & (e1 = op(e0, e1)) & (e0 = op(e0, e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG094+1.p', ax2)).
fof(f1073, plain, ((e2 = op(e2, e1)) | ~ sP0), inference(subsumption_resolution, [], [f1072, f182])).
fof(f182, plain, (e1 = op(e2, e2)), inference(cnf_transformation, [], [f2])).
fof(f1072, plain, ((e2 = op(e2, e1)) | ~ (e1 = op(e2, e2)) | ~ sP0), inference(subsumption_resolution, [], [f582, f176])).
fof(f176, plain, (e2 = op(e1, e1)), inference(cnf_transformation, [], [f2])).
fof(f582, plain, ((e2 = op(e2, e1)) | ~ (e2 = op(e1, e1)) | ~ (e1 = op(e2, e2)) | ~ sP0), inference(cnf_transformation, [], [f159])).
fof(f159, plain, ((((e4 = op(e4, e4)) | ~ (e4 = op(e4, e4)) | ~ (e4 = op(e4, e4))) & ((e4 = op(e4, e3)) | ~ (e4 = op(e3, e3)) | ~ (e3 = op(e4, e4))) & ((e4 = op(e4, e2)) | ~ (e4 = op(e2, e2)) | ~ (e2 = op(e4, e4))) & ((e4 = op(e4, e1)) | ~ (e4 = op(e1, e1)) | ~ (e1 = op(e4, e4))) & ((e4 = op(e4, e0)) | ~ (e4 = op(e0, e0)) | ~ (e0 = op(e4, e4))) & ((e3 = op(e3, e4)) | ~ (e3 = op(e4, e4)) | ~ (e4 = op(e3, e3))) & ((e3 = op(e3, e3)) | ~ (e3 = op(e3, e3)) | ~ (e3 = op(e3, e3))) & ((e3 = op(e3, e2)) | ~ (e3 = op(e2, e2)) | ~ (e2 = op(e3, e3))) & ((e3 = op(e3, e1)) | ~ (e3 = op(e1, e1)) | ~ (e1 = op(e3, e3))) & ((e3 = op(e3, e0)) | ~ (e3 = op(e0, e0)) | ~ (e0 = op(e3, e3))) & ((e2 = op(e2, e4)) | ~ (e2 = op(e4, e4)) | ~ (e4 = op(e2, e2))) & ((e2 = op(e2, e3)) | ~ (e2 = op(e3, e3)) | ~ (e3 = op(e2, e2))) & ((e2 = op(e2, e2)) | ~ (e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & ((e2 = op(e2, e1)) | ~ (e2 = op(e1, e1)) | ~ (e1 = op(e2, e2))) & ((e2 = op(e2, e0)) | ~ (e2 = op(e0, e0)) | ~ (e0 = op(e2, e2))) & ((e1 = op(e1, e4)) | ~ (e1 = op(e4, e4)) | ~ (e4 = op(e1, e1))) & ((e1 = op(e1, e3)) | ~ (e1 = op(e3, e3)) | ~ (e3 = op(e1, e1))) & ((e1 = op(e1, e2)) | ~ (e1 = op(e2, e2)) | ~ (e2 = op(e1, e1))) & ((e1 = op(e1, e1)) | ~ (e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & ((e1 = op(e1, e0)) | ~ (e1 = op(e0, e0)) | ~ (e0 = op(e1, e1))) & ((e0 = op(e0, e4)) | ~ (e0 = op(e4, e4)) | ~ (e4 = op(e0, e0))) & ((e0 = op(e0, e3)) | ~ (e0 = op(e3, e3)) | ~ (e3 = op(e0, e0))) & ((e0 = op(e0, e2)) | ~ (e0 = op(e2, e2)) | ~ (e2 = op(e0, e0))) & ((e0 = op(e0, e1)) | ~ (e0 = op(e1, e1)) | ~ (e1 = op(e0, e0))) & ((e0 = op(e0, e0)) | ~ (e0 = op(e0, e0)) | ~ (e0 = op(e0, e0)))) | ~ sP0), inference(nnf_transformation, [], [f7])).
fof(f7, plain, ((((e4 = op(e4, e4)) | ~ (e4 = op(e4, e4)) | ~ (e4 = op(e4, e4))) & ((e4 = op(e4, e3)) | ~ (e4 = op(e3, e3)) | ~ (e3 = op(e4, e4))) & ((e4 = op(e4, e2)) | ~ (e4 = op(e2, e2)) | ~ (e2 = op(e4, e4))) & ((e4 = op(e4, e1)) | ~ (e4 = op(e1, e1)) | ~ (e1 = op(e4, e4))) & ((e4 = op(e4, e0)) | ~ (e4 = op(e0, e0)) | ~ (e0 = op(e4, e4))) & ((e3 = op(e3, e4)) | ~ (e3 = op(e4, e4)) | ~ (e4 = op(e3, e3))) & ((e3 = op(e3, e3)) | ~ (e3 = op(e3, e3)) | ~ (e3 = op(e3, e3))) & ((e3 = op(e3, e2)) | ~ (e3 = op(e2, e2)) | ~ (e2 = op(e3, e3))) & ((e3 = op(e3, e1)) | ~ (e3 = op(e1, e1)) | ~ (e1 = op(e3, e3))) & ((e3 = op(e3, e0)) | ~ (e3 = op(e0, e0)) | ~ (e0 = op(e3, e3))) & ((e2 = op(e2, e4)) | ~ (e2 = op(e4, e4)) | ~ (e4 = op(e2, e2))) & ((e2 = op(e2, e3)) | ~ (e2 = op(e3, e3)) | ~ (e3 = op(e2, e2))) & ((e2 = op(e2, e2)) | ~ (e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & ((e2 = op(e2, e1)) | ~ (e2 = op(e1, e1)) | ~ (e1 = op(e2, e2))) & ((e2 = op(e2, e0)) | ~ (e2 = op(e0, e0)) | ~ (e0 = op(e2, e2))) & ((e1 = op(e1, e4)) | ~ (e1 = op(e4, e4)) | ~ (e4 = op(e1, e1))) & ((e1 = op(e1, e3)) | ~ (e1 = op(e3, e3)) | ~ (e3 = op(e1, e1))) & ((e1 = op(e1, e2)) | ~ (e1 = op(e2, e2)) | ~ (e2 = op(e1, e1))) & ((e1 = op(e1, e1)) | ~ (e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & ((e1 = op(e1, e0)) | ~ (e1 = op(e0, e0)) | ~ (e0 = op(e1, e1))) & ((e0 = op(e0, e4)) | ~ (e0 = op(e4, e4)) | ~ (e4 = op(e0, e0))) & ((e0 = op(e0, e3)) | ~ (e0 = op(e3, e3)) | ~ (e3 = op(e0, e0))) & ((e0 = op(e0, e2)) | ~ (e0 = op(e2, e2)) | ~ (e2 = op(e0, e0))) & ((e0 = op(e0, e1)) | ~ (e0 = op(e1, e1)) | ~ (e1 = op(e0, e0))) & ((e0 = op(e0, e0)) | ~ (e0 = op(e0, e0)) | ~ (e0 = op(e0, e0)))) | ~ sP0), inference(usedef, [], [e7])).
fof(e7, plain, (sP0 <=> (((e4 = op(e4, e4)) | ~ (e4 = op(e4, e4)) | ~ (e4 = op(e4, e4))) & ((e4 = op(e4, e3)) | ~ (e4 = op(e3, e3)) | ~ (e3 = op(e4, e4))) & ((e4 = op(e4, e2)) | ~ (e4 = op(e2, e2)) | ~ (e2 = op(e4, e4))) & ((e4 = op(e4, e1)) | ~ (e4 = op(e1, e1)) | ~ (e1 = op(e4, e4))) & ((e4 = op(e4, e0)) | ~ (e4 = op(e0, e0)) | ~ (e0 = op(e4, e4))) & ((e3 = op(e3, e4)) | ~ (e3 = op(e4, e4)) | ~ (e4 = op(e3, e3))) & ((e3 = op(e3, e3)) | ~ (e3 = op(e3, e3)) | ~ (e3 = op(e3, e3))) & ((e3 = op(e3, e2)) | ~ (e3 = op(e2, e2)) | ~ (e2 = op(e3, e3))) & ((e3 = op(e3, e1)) | ~ (e3 = op(e1, e1)) | ~ (e1 = op(e3, e3))) & ((e3 = op(e3, e0)) | ~ (e3 = op(e0, e0)) | ~ (e0 = op(e3, e3))) & ((e2 = op(e2, e4)) | ~ (e2 = op(e4, e4)) | ~ (e4 = op(e2, e2))) & ((e2 = op(e2, e3)) | ~ (e2 = op(e3, e3)) | ~ (e3 = op(e2, e2))) & ((e2 = op(e2, e2)) | ~ (e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & ((e2 = op(e2, e1)) | ~ (e2 = op(e1, e1)) | ~ (e1 = op(e2, e2))) & ((e2 = op(e2, e0)) | ~ (e2 = op(e0, e0)) | ~ (e0 = op(e2, e2))) & ((e1 = op(e1, e4)) | ~ (e1 = op(e4, e4)) | ~ (e4 = op(e1, e1))) & ((e1 = op(e1, e3)) | ~ (e1 = op(e3, e3)) | ~ (e3 = op(e1, e1))) & ((e1 = op(e1, e2)) | ~ (e1 = op(e2, e2)) | ~ (e2 = op(e1, e1))) & ((e1 = op(e1, e1)) | ~ (e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & ((e1 = op(e1, e0)) | ~ (e1 = op(e0, e0)) | ~ (e0 = op(e1, e1))) & ((e0 = op(e0, e4)) | ~ (e0 = op(e4, e4)) | ~ (e4 = op(e0, e0))) & ((e0 = op(e0, e3)) | ~ (e0 = op(e3, e3)) | ~ (e3 = op(e0, e0))) & ((e0 = op(e0, e2)) | ~ (e0 = op(e2, e2)) | ~ (e2 = op(e0, e0))) & ((e0 = op(e0, e1)) | ~ (e0 = op(e1, e1)) | ~ (e1 = op(e0, e0))) & ((e0 = op(e0, e0)) | ~ (e0 = op(e0, e0)) | ~ (e0 = op(e0, e0))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f1502, plain, sP0, inference(subsumption_resolution, [], [f1501, f1045])).
fof(f1045, plain, ~ sP1, inference(subsumption_resolution, [], [f800, f617])).
fof(f617, plain, (unit = op(unit, unit)), inference(definition_unfolding, [], [f170, f195, f195, f195])).
fof(f195, plain, (e0 = unit), inference(cnf_transformation, [], [f3])).
fof(f3, plain, (e0 = unit), file('/home/ubuntu/library/tptp/Problems/ALG/ALG094+1.p', ax3)).
fof(f170, plain, (e0 = op(e0, e0)), inference(cnf_transformation, [], [f2])).
fof(f800, plain, (~ (unit = op(unit, unit)) | ~ sP1), inference(definition_unfolding, [], [f566, f195, f195, f195])).
fof(f566, plain, (~ (e0 = op(e0, e0)) | ~ sP1), inference(cnf_transformation, [], [f158])).
fof(f158, plain, ((~ (e4 = op(e0, e0)) & ~ (e3 = op(e0, e0)) & ~ (e2 = op(e0, e0)) & ~ (e1 = op(e0, e0)) & ~ (e0 = op(e0, e0))) | ~ sP1), inference(nnf_transformation, [], [f8])).
fof(f8, plain, ((~ (e4 = op(e0, e0)) & ~ (e3 = op(e0, e0)) & ~ (e2 = op(e0, e0)) & ~ (e1 = op(e0, e0)) & ~ (e0 = op(e0, e0))) | ~ sP1), inference(usedef, [], [e8])).
fof(e8, plain, (sP1 <=> (~ (e4 = op(e0, e0)) & ~ (e3 = op(e0, e0)) & ~ (e2 = op(e0, e0)) & ~ (e1 = op(e0, e0)) & ~ (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f1501, plain, (sP1 | sP0), inference(subsumption_resolution, [], [f1500, f1040])).
fof(f1040, plain, ~ sP2, inference(subsumption_resolution, [], [f794, f616])).
fof(f616, plain, (e1 = op(unit, e1)), inference(definition_unfolding, [], [f171, f195])).
fof(f171, plain, (e1 = op(e0, e1)), inference(cnf_transformation, [], [f2])).
fof(f794, plain, (~ (e1 = op(unit, e1)) | ~ sP2), inference(definition_unfolding, [], [f562, f195])).
fof(f562, plain, (~ (e1 = op(e0, e1)) | ~ sP2), inference(cnf_transformation, [], [f157])).
fof(f157, plain, ((~ (e4 = op(e0, e1)) & ~ (e3 = op(e0, e1)) & ~ (e2 = op(e0, e1)) & ~ (e1 = op(e0, e1)) & ~ (e0 = op(e0, e1))) | ~ sP2), inference(nnf_transformation, [], [f9])).
fof(f9, plain, ((~ (e4 = op(e0, e1)) & ~ (e3 = op(e0, e1)) & ~ (e2 = op(e0, e1)) & ~ (e1 = op(e0, e1)) & ~ (e0 = op(e0, e1))) | ~ sP2), inference(usedef, [], [e9])).
fof(e9, plain, (sP2 <=> (~ (e4 = op(e0, e1)) & ~ (e3 = op(e0, e1)) & ~ (e2 = op(e0, e1)) & ~ (e1 = op(e0, e1)) & ~ (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f1500, plain, (sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1499, f1036])).
fof(f1036, plain, ~ sP3, inference(subsumption_resolution, [], [f788, f615])).
fof(f615, plain, (e2 = op(unit, e2)), inference(definition_unfolding, [], [f172, f195])).
fof(f172, plain, (e2 = op(e0, e2)), inference(cnf_transformation, [], [f2])).
fof(f788, plain, (~ (e2 = op(unit, e2)) | ~ sP3), inference(definition_unfolding, [], [f558, f195])).
fof(f558, plain, (~ (e2 = op(e0, e2)) | ~ sP3), inference(cnf_transformation, [], [f156])).
fof(f156, plain, ((~ (e4 = op(e0, e2)) & ~ (e3 = op(e0, e2)) & ~ (e2 = op(e0, e2)) & ~ (e1 = op(e0, e2)) & ~ (e0 = op(e0, e2))) | ~ sP3), inference(nnf_transformation, [], [f10])).
fof(f10, plain, ((~ (e4 = op(e0, e2)) & ~ (e3 = op(e0, e2)) & ~ (e2 = op(e0, e2)) & ~ (e1 = op(e0, e2)) & ~ (e0 = op(e0, e2))) | ~ sP3), inference(usedef, [], [e10])).
fof(e10, plain, (sP3 <=> (~ (e4 = op(e0, e2)) & ~ (e3 = op(e0, e2)) & ~ (e2 = op(e0, e2)) & ~ (e1 = op(e0, e2)) & ~ (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f1499, plain, (sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1498, f1033])).
fof(f1033, plain, ~ sP4, inference(subsumption_resolution, [], [f782, f614])).
fof(f614, plain, (e3 = op(unit, e3)), inference(definition_unfolding, [], [f173, f195])).
fof(f173, plain, (e3 = op(e0, e3)), inference(cnf_transformation, [], [f2])).
fof(f782, plain, (~ (e3 = op(unit, e3)) | ~ sP4), inference(definition_unfolding, [], [f554, f195])).
fof(f554, plain, (~ (e3 = op(e0, e3)) | ~ sP4), inference(cnf_transformation, [], [f155])).
fof(f155, plain, ((~ (e4 = op(e0, e3)) & ~ (e3 = op(e0, e3)) & ~ (e2 = op(e0, e3)) & ~ (e1 = op(e0, e3)) & ~ (e0 = op(e0, e3))) | ~ sP4), inference(nnf_transformation, [], [f11])).
fof(f11, plain, ((~ (e4 = op(e0, e3)) & ~ (e3 = op(e0, e3)) & ~ (e2 = op(e0, e3)) & ~ (e1 = op(e0, e3)) & ~ (e0 = op(e0, e3))) | ~ sP4), inference(usedef, [], [e11])).
fof(e11, plain, (sP4 <=> (~ (e4 = op(e0, e3)) & ~ (e3 = op(e0, e3)) & ~ (e2 = op(e0, e3)) & ~ (e1 = op(e0, e3)) & ~ (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f1498, plain, (sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1497, f1031])).
fof(f1031, plain, ~ sP5, inference(subsumption_resolution, [], [f776, f613])).
fof(f613, plain, (e4 = op(unit, e4)), inference(definition_unfolding, [], [f174, f195])).
fof(f174, plain, (e4 = op(e0, e4)), inference(cnf_transformation, [], [f2])).
fof(f776, plain, (~ (e4 = op(unit, e4)) | ~ sP5), inference(definition_unfolding, [], [f550, f195])).
fof(f550, plain, (~ (e4 = op(e0, e4)) | ~ sP5), inference(cnf_transformation, [], [f154])).
fof(f154, plain, ((~ (e4 = op(e0, e4)) & ~ (e3 = op(e0, e4)) & ~ (e2 = op(e0, e4)) & ~ (e1 = op(e0, e4)) & ~ (e0 = op(e0, e4))) | ~ sP5), inference(nnf_transformation, [], [f12])).
fof(f12, plain, ((~ (e4 = op(e0, e4)) & ~ (e3 = op(e0, e4)) & ~ (e2 = op(e0, e4)) & ~ (e1 = op(e0, e4)) & ~ (e0 = op(e0, e4))) | ~ sP5), inference(usedef, [], [e12])).
fof(e12, plain, (sP5 <=> (~ (e4 = op(e0, e4)) & ~ (e3 = op(e0, e4)) & ~ (e2 = op(e0, e4)) & ~ (e1 = op(e0, e4)) & ~ (e0 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f1497, plain, (sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1496, f1030])).
fof(f1030, plain, ~ sP6, inference(subsumption_resolution, [], [f774, f612])).
fof(f612, plain, (e1 = op(e1, unit)), inference(definition_unfolding, [], [f175, f195])).
fof(f175, plain, (e1 = op(e1, e0)), inference(cnf_transformation, [], [f2])).
fof(f774, plain, (~ (e1 = op(e1, unit)) | ~ sP6), inference(definition_unfolding, [], [f542, f195])).
fof(f542, plain, (~ (e1 = op(e1, e0)) | ~ sP6), inference(cnf_transformation, [], [f153])).
fof(f153, plain, ((~ (e4 = op(e1, e0)) & ~ (e3 = op(e1, e0)) & ~ (e2 = op(e1, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e1, e0))) | ~ sP6), inference(nnf_transformation, [], [f13])).
fof(f13, plain, ((~ (e4 = op(e1, e0)) & ~ (e3 = op(e1, e0)) & ~ (e2 = op(e1, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e1, e0))) | ~ sP6), inference(usedef, [], [e13])).
fof(e13, plain, (sP6 <=> (~ (e4 = op(e1, e0)) & ~ (e3 = op(e1, e0)) & ~ (e2 = op(e1, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f1496, plain, (sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1495, f1026])).
fof(f1026, plain, ~ sP7, inference(subsumption_resolution, [], [f538, f176])).
fof(f538, plain, (~ (e2 = op(e1, e1)) | ~ sP7), inference(cnf_transformation, [], [f152])).
fof(f152, plain, ((~ (e4 = op(e1, e1)) & ~ (e3 = op(e1, e1)) & ~ (e2 = op(e1, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e1, e1))) | ~ sP7), inference(nnf_transformation, [], [f14])).
fof(f14, plain, ((~ (e4 = op(e1, e1)) & ~ (e3 = op(e1, e1)) & ~ (e2 = op(e1, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e1, e1))) | ~ sP7), inference(usedef, [], [e14])).
fof(e14, plain, (sP7 <=> (~ (e4 = op(e1, e1)) & ~ (e3 = op(e1, e1)) & ~ (e2 = op(e1, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f1495, plain, (sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1494, f1023])).
fof(f1023, plain, ~ sP8, inference(subsumption_resolution, [], [f534, f177])).
fof(f177, plain, (e3 = op(e1, e2)), inference(cnf_transformation, [], [f2])).
fof(f534, plain, (~ (e3 = op(e1, e2)) | ~ sP8), inference(cnf_transformation, [], [f151])).
fof(f151, plain, ((~ (e4 = op(e1, e2)) & ~ (e3 = op(e1, e2)) & ~ (e2 = op(e1, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e1, e2))) | ~ sP8), inference(nnf_transformation, [], [f15])).
fof(f15, plain, ((~ (e4 = op(e1, e2)) & ~ (e3 = op(e1, e2)) & ~ (e2 = op(e1, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e1, e2))) | ~ sP8), inference(usedef, [], [e15])).
fof(e15, plain, (sP8 <=> (~ (e4 = op(e1, e2)) & ~ (e3 = op(e1, e2)) & ~ (e2 = op(e1, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f1494, plain, (sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1493, f1021])).
fof(f1021, plain, ~ sP9, inference(subsumption_resolution, [], [f530, f178])).
fof(f178, plain, (e4 = op(e1, e3)), inference(cnf_transformation, [], [f2])).
fof(f530, plain, (~ (e4 = op(e1, e3)) | ~ sP9), inference(cnf_transformation, [], [f150])).
fof(f150, plain, ((~ (e4 = op(e1, e3)) & ~ (e3 = op(e1, e3)) & ~ (e2 = op(e1, e3)) & ~ (e1 = op(e1, e3)) & ~ (e0 = op(e1, e3))) | ~ sP9), inference(nnf_transformation, [], [f16])).
fof(f16, plain, ((~ (e4 = op(e1, e3)) & ~ (e3 = op(e1, e3)) & ~ (e2 = op(e1, e3)) & ~ (e1 = op(e1, e3)) & ~ (e0 = op(e1, e3))) | ~ sP9), inference(usedef, [], [e16])).
fof(e16, plain, (sP9 <=> (~ (e4 = op(e1, e3)) & ~ (e3 = op(e1, e3)) & ~ (e2 = op(e1, e3)) & ~ (e1 = op(e1, e3)) & ~ (e0 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f1493, plain, (sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1492, f1020])).
fof(f1020, plain, ~ sP10, inference(subsumption_resolution, [], [f767, f611])).
fof(f611, plain, (op(e1, e4) = unit), inference(definition_unfolding, [], [f179, f195])).
fof(f179, plain, (e0 = op(e1, e4)), inference(cnf_transformation, [], [f2])).
fof(f767, plain, (~ (op(e1, e4) = unit) | ~ sP10), inference(definition_unfolding, [], [f521, f195])).
fof(f521, plain, (~ (e0 = op(e1, e4)) | ~ sP10), inference(cnf_transformation, [], [f149])).
fof(f149, plain, ((~ (e4 = op(e1, e4)) & ~ (e3 = op(e1, e4)) & ~ (e2 = op(e1, e4)) & ~ (e1 = op(e1, e4)) & ~ (e0 = op(e1, e4))) | ~ sP10), inference(nnf_transformation, [], [f17])).
fof(f17, plain, ((~ (e4 = op(e1, e4)) & ~ (e3 = op(e1, e4)) & ~ (e2 = op(e1, e4)) & ~ (e1 = op(e1, e4)) & ~ (e0 = op(e1, e4))) | ~ sP10), inference(usedef, [], [e17])).
fof(e17, plain, (sP10 <=> (~ (e4 = op(e1, e4)) & ~ (e3 = op(e1, e4)) & ~ (e2 = op(e1, e4)) & ~ (e1 = op(e1, e4)) & ~ (e0 = op(e1, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f1492, plain, (sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1491, f1015])).
fof(f1015, plain, ~ sP11, inference(subsumption_resolution, [], [f764, f610])).
fof(f610, plain, (e2 = op(e2, unit)), inference(definition_unfolding, [], [f180, f195])).
fof(f180, plain, (e2 = op(e2, e0)), inference(cnf_transformation, [], [f2])).
fof(f764, plain, (~ (e2 = op(e2, unit)) | ~ sP11), inference(definition_unfolding, [], [f518, f195])).
fof(f518, plain, (~ (e2 = op(e2, e0)) | ~ sP11), inference(cnf_transformation, [], [f148])).
fof(f148, plain, ((~ (e4 = op(e2, e0)) & ~ (e3 = op(e2, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e2, e0)) & ~ (e0 = op(e2, e0))) | ~ sP11), inference(nnf_transformation, [], [f18])).
fof(f18, plain, ((~ (e4 = op(e2, e0)) & ~ (e3 = op(e2, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e2, e0)) & ~ (e0 = op(e2, e0))) | ~ sP11), inference(usedef, [], [e18])).
fof(e18, plain, (sP11 <=> (~ (e4 = op(e2, e0)) & ~ (e3 = op(e2, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e2, e0)) & ~ (e0 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f1491, plain, (sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1490, f1012])).
fof(f1012, plain, ~ sP12, inference(subsumption_resolution, [], [f515, f181])).
fof(f515, plain, (~ (e4 = op(e2, e1)) | ~ sP12), inference(cnf_transformation, [], [f147])).
fof(f147, plain, ((~ (e4 = op(e2, e1)) & ~ (e3 = op(e2, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e2, e1)) & ~ (e0 = op(e2, e1))) | ~ sP12), inference(nnf_transformation, [], [f19])).
fof(f19, plain, ((~ (e4 = op(e2, e1)) & ~ (e3 = op(e2, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e2, e1)) & ~ (e0 = op(e2, e1))) | ~ sP12), inference(usedef, [], [e19])).
fof(e19, plain, (sP12 <=> (~ (e4 = op(e2, e1)) & ~ (e3 = op(e2, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e2, e1)) & ~ (e0 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f1490, plain, (sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1489, f1011])).
fof(f1011, plain, ~ sP13, inference(subsumption_resolution, [], [f507, f182])).
fof(f507, plain, (~ (e1 = op(e2, e2)) | ~ sP13), inference(cnf_transformation, [], [f146])).
fof(f146, plain, ((~ (e4 = op(e2, e2)) & ~ (e3 = op(e2, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e2, e2)) & ~ (e0 = op(e2, e2))) | ~ sP13), inference(nnf_transformation, [], [f20])).
fof(f20, plain, ((~ (e4 = op(e2, e2)) & ~ (e3 = op(e2, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e2, e2)) & ~ (e0 = op(e2, e2))) | ~ sP13), inference(usedef, [], [e20])).
fof(e20, plain, (sP13 <=> (~ (e4 = op(e2, e2)) & ~ (e3 = op(e2, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e2, e2)) & ~ (e0 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f1489, plain, (sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1488, f1007])).
fof(f1007, plain, ~ sP14, inference(subsumption_resolution, [], [f759, f609])).
fof(f609, plain, (op(e2, e3) = unit), inference(definition_unfolding, [], [f183, f195])).
fof(f183, plain, (e0 = op(e2, e3)), inference(cnf_transformation, [], [f2])).
fof(f759, plain, (~ (op(e2, e3) = unit) | ~ sP14), inference(definition_unfolding, [], [f501, f195])).
fof(f501, plain, (~ (e0 = op(e2, e3)) | ~ sP14), inference(cnf_transformation, [], [f145])).
fof(f145, plain, ((~ (e4 = op(e2, e3)) & ~ (e3 = op(e2, e3)) & ~ (e2 = op(e2, e3)) & ~ (e1 = op(e2, e3)) & ~ (e0 = op(e2, e3))) | ~ sP14), inference(nnf_transformation, [], [f21])).
fof(f21, plain, ((~ (e4 = op(e2, e3)) & ~ (e3 = op(e2, e3)) & ~ (e2 = op(e2, e3)) & ~ (e1 = op(e2, e3)) & ~ (e0 = op(e2, e3))) | ~ sP14), inference(usedef, [], [e21])).
fof(e21, plain, (sP14 <=> (~ (e4 = op(e2, e3)) & ~ (e3 = op(e2, e3)) & ~ (e2 = op(e2, e3)) & ~ (e1 = op(e2, e3)) & ~ (e0 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f1488, plain, (sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1487, f1002])).
fof(f1002, plain, ~ sP15, inference(subsumption_resolution, [], [f499, f184])).
fof(f184, plain, (e3 = op(e2, e4)), inference(cnf_transformation, [], [f2])).
fof(f499, plain, (~ (e3 = op(e2, e4)) | ~ sP15), inference(cnf_transformation, [], [f144])).
fof(f144, plain, ((~ (e4 = op(e2, e4)) & ~ (e3 = op(e2, e4)) & ~ (e2 = op(e2, e4)) & ~ (e1 = op(e2, e4)) & ~ (e0 = op(e2, e4))) | ~ sP15), inference(nnf_transformation, [], [f22])).
fof(f22, plain, ((~ (e4 = op(e2, e4)) & ~ (e3 = op(e2, e4)) & ~ (e2 = op(e2, e4)) & ~ (e1 = op(e2, e4)) & ~ (e0 = op(e2, e4))) | ~ sP15), inference(usedef, [], [e22])).
fof(e22, plain, (sP15 <=> (~ (e4 = op(e2, e4)) & ~ (e3 = op(e2, e4)) & ~ (e2 = op(e2, e4)) & ~ (e1 = op(e2, e4)) & ~ (e0 = op(e2, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f1487, plain, (sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1486, f1000])).
fof(f1000, plain, ~ sP16, inference(subsumption_resolution, [], [f754, f608])).
fof(f608, plain, (e3 = op(e3, unit)), inference(definition_unfolding, [], [f185, f195])).
fof(f185, plain, (e3 = op(e3, e0)), inference(cnf_transformation, [], [f2])).
fof(f754, plain, (~ (e3 = op(e3, unit)) | ~ sP16), inference(definition_unfolding, [], [f494, f195])).
fof(f494, plain, (~ (e3 = op(e3, e0)) | ~ sP16), inference(cnf_transformation, [], [f143])).
fof(f143, plain, ((~ (e4 = op(e3, e0)) & ~ (e3 = op(e3, e0)) & ~ (e2 = op(e3, e0)) & ~ (e1 = op(e3, e0)) & ~ (e0 = op(e3, e0))) | ~ sP16), inference(nnf_transformation, [], [f23])).
fof(f23, plain, ((~ (e4 = op(e3, e0)) & ~ (e3 = op(e3, e0)) & ~ (e2 = op(e3, e0)) & ~ (e1 = op(e3, e0)) & ~ (e0 = op(e3, e0))) | ~ sP16), inference(usedef, [], [e23])).
fof(e23, plain, (sP16 <=> (~ (e4 = op(e3, e0)) & ~ (e3 = op(e3, e0)) & ~ (e2 = op(e3, e0)) & ~ (e1 = op(e3, e0)) & ~ (e0 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f1486, plain, (sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1485, f998])).
fof(f998, plain, ~ sP17, inference(subsumption_resolution, [], [f752, f607])).
fof(f607, plain, (op(e3, e1) = unit), inference(definition_unfolding, [], [f186, f195])).
fof(f186, plain, (e0 = op(e3, e1)), inference(cnf_transformation, [], [f2])).
fof(f752, plain, (~ (op(e3, e1) = unit) | ~ sP17), inference(definition_unfolding, [], [f486, f195])).
fof(f486, plain, (~ (e0 = op(e3, e1)) | ~ sP17), inference(cnf_transformation, [], [f142])).
fof(f142, plain, ((~ (e4 = op(e3, e1)) & ~ (e3 = op(e3, e1)) & ~ (e2 = op(e3, e1)) & ~ (e1 = op(e3, e1)) & ~ (e0 = op(e3, e1))) | ~ sP17), inference(nnf_transformation, [], [f24])).
fof(f24, plain, ((~ (e4 = op(e3, e1)) & ~ (e3 = op(e3, e1)) & ~ (e2 = op(e3, e1)) & ~ (e1 = op(e3, e1)) & ~ (e0 = op(e3, e1))) | ~ sP17), inference(usedef, [], [e24])).
fof(e24, plain, (sP17 <=> (~ (e4 = op(e3, e1)) & ~ (e3 = op(e3, e1)) & ~ (e2 = op(e3, e1)) & ~ (e1 = op(e3, e1)) & ~ (e0 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f1485, plain, (sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1484, f993])).
fof(f993, plain, ~ sP18, inference(subsumption_resolution, [], [f485, f187])).
fof(f187, plain, (e4 = op(e3, e2)), inference(cnf_transformation, [], [f2])).
fof(f485, plain, (~ (e4 = op(e3, e2)) | ~ sP18), inference(cnf_transformation, [], [f141])).
fof(f141, plain, ((~ (e4 = op(e3, e2)) & ~ (e3 = op(e3, e2)) & ~ (e2 = op(e3, e2)) & ~ (e1 = op(e3, e2)) & ~ (e0 = op(e3, e2))) | ~ sP18), inference(nnf_transformation, [], [f25])).
fof(f25, plain, ((~ (e4 = op(e3, e2)) & ~ (e3 = op(e3, e2)) & ~ (e2 = op(e3, e2)) & ~ (e1 = op(e3, e2)) & ~ (e0 = op(e3, e2))) | ~ sP18), inference(usedef, [], [e25])).
fof(e25, plain, (sP18 <=> (~ (e4 = op(e3, e2)) & ~ (e3 = op(e3, e2)) & ~ (e2 = op(e3, e2)) & ~ (e1 = op(e3, e2)) & ~ (e0 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f1484, plain, (sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1483, f992])).
fof(f992, plain, ~ sP19, inference(subsumption_resolution, [], [f477, f188])).
fof(f188, plain, (e1 = op(e3, e3)), inference(cnf_transformation, [], [f2])).
fof(f477, plain, (~ (e1 = op(e3, e3)) | ~ sP19), inference(cnf_transformation, [], [f140])).
fof(f140, plain, ((~ (e4 = op(e3, e3)) & ~ (e3 = op(e3, e3)) & ~ (e2 = op(e3, e3)) & ~ (e1 = op(e3, e3)) & ~ (e0 = op(e3, e3))) | ~ sP19), inference(nnf_transformation, [], [f26])).
fof(f26, plain, ((~ (e4 = op(e3, e3)) & ~ (e3 = op(e3, e3)) & ~ (e2 = op(e3, e3)) & ~ (e1 = op(e3, e3)) & ~ (e0 = op(e3, e3))) | ~ sP19), inference(usedef, [], [e26])).
fof(e26, plain, (sP19 <=> (~ (e4 = op(e3, e3)) & ~ (e3 = op(e3, e3)) & ~ (e2 = op(e3, e3)) & ~ (e1 = op(e3, e3)) & ~ (e0 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP19])])).
fof(f1483, plain, (sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1482, f988])).
fof(f988, plain, ~ sP20, inference(subsumption_resolution, [], [f473, f189])).
fof(f189, plain, (e2 = op(e3, e4)), inference(cnf_transformation, [], [f2])).
fof(f473, plain, (~ (e2 = op(e3, e4)) | ~ sP20), inference(cnf_transformation, [], [f139])).
fof(f139, plain, ((~ (e4 = op(e3, e4)) & ~ (e3 = op(e3, e4)) & ~ (e2 = op(e3, e4)) & ~ (e1 = op(e3, e4)) & ~ (e0 = op(e3, e4))) | ~ sP20), inference(nnf_transformation, [], [f27])).
fof(f27, plain, ((~ (e4 = op(e3, e4)) & ~ (e3 = op(e3, e4)) & ~ (e2 = op(e3, e4)) & ~ (e1 = op(e3, e4)) & ~ (e0 = op(e3, e4))) | ~ sP20), inference(usedef, [], [e27])).
fof(e27, plain, (sP20 <=> (~ (e4 = op(e3, e4)) & ~ (e3 = op(e3, e4)) & ~ (e2 = op(e3, e4)) & ~ (e1 = op(e3, e4)) & ~ (e0 = op(e3, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP20])])).
fof(f1482, plain, (sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1481, f985])).
fof(f985, plain, ~ sP21, inference(subsumption_resolution, [], [f744, f606])).
fof(f606, plain, (e4 = op(e4, unit)), inference(definition_unfolding, [], [f190, f195])).
fof(f190, plain, (e4 = op(e4, e0)), inference(cnf_transformation, [], [f2])).
fof(f744, plain, (~ (e4 = op(e4, unit)) | ~ sP21), inference(definition_unfolding, [], [f470, f195])).
fof(f470, plain, (~ (e4 = op(e4, e0)) | ~ sP21), inference(cnf_transformation, [], [f138])).
fof(f138, plain, ((~ (e4 = op(e4, e0)) & ~ (e3 = op(e4, e0)) & ~ (e2 = op(e4, e0)) & ~ (e1 = op(e4, e0)) & ~ (e0 = op(e4, e0))) | ~ sP21), inference(nnf_transformation, [], [f28])).
fof(f28, plain, ((~ (e4 = op(e4, e0)) & ~ (e3 = op(e4, e0)) & ~ (e2 = op(e4, e0)) & ~ (e1 = op(e4, e0)) & ~ (e0 = op(e4, e0))) | ~ sP21), inference(usedef, [], [e28])).
fof(e28, plain, (sP21 <=> (~ (e4 = op(e4, e0)) & ~ (e3 = op(e4, e0)) & ~ (e2 = op(e4, e0)) & ~ (e1 = op(e4, e0)) & ~ (e0 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP21])])).
fof(f1481, plain, (sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1480, f984])).
fof(f984, plain, ~ sP22, inference(subsumption_resolution, [], [f464, f191])).
fof(f191, plain, (e3 = op(e4, e1)), inference(cnf_transformation, [], [f2])).
fof(f464, plain, (~ (e3 = op(e4, e1)) | ~ sP22), inference(cnf_transformation, [], [f137])).
fof(f137, plain, ((~ (e4 = op(e4, e1)) & ~ (e3 = op(e4, e1)) & ~ (e2 = op(e4, e1)) & ~ (e1 = op(e4, e1)) & ~ (e0 = op(e4, e1))) | ~ sP22), inference(nnf_transformation, [], [f29])).
fof(f29, plain, ((~ (e4 = op(e4, e1)) & ~ (e3 = op(e4, e1)) & ~ (e2 = op(e4, e1)) & ~ (e1 = op(e4, e1)) & ~ (e0 = op(e4, e1))) | ~ sP22), inference(usedef, [], [e29])).
fof(e29, plain, (sP22 <=> (~ (e4 = op(e4, e1)) & ~ (e3 = op(e4, e1)) & ~ (e2 = op(e4, e1)) & ~ (e1 = op(e4, e1)) & ~ (e0 = op(e4, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP22])])).
fof(f1480, plain, (sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1479, f982])).
fof(f982, plain, ~ sP23, inference(subsumption_resolution, [], [f742, f605])).
fof(f605, plain, (op(e4, e2) = unit), inference(definition_unfolding, [], [f192, f195])).
fof(f192, plain, (e0 = op(e4, e2)), inference(cnf_transformation, [], [f2])).
fof(f742, plain, (~ (op(e4, e2) = unit) | ~ sP23), inference(definition_unfolding, [], [f456, f195])).
fof(f456, plain, (~ (e0 = op(e4, e2)) | ~ sP23), inference(cnf_transformation, [], [f136])).
fof(f136, plain, ((~ (e4 = op(e4, e2)) & ~ (e3 = op(e4, e2)) & ~ (e2 = op(e4, e2)) & ~ (e1 = op(e4, e2)) & ~ (e0 = op(e4, e2))) | ~ sP23), inference(nnf_transformation, [], [f30])).
fof(f30, plain, ((~ (e4 = op(e4, e2)) & ~ (e3 = op(e4, e2)) & ~ (e2 = op(e4, e2)) & ~ (e1 = op(e4, e2)) & ~ (e0 = op(e4, e2))) | ~ sP23), inference(usedef, [], [e30])).
fof(e30, plain, (sP23 <=> (~ (e4 = op(e4, e2)) & ~ (e3 = op(e4, e2)) & ~ (e2 = op(e4, e2)) & ~ (e1 = op(e4, e2)) & ~ (e0 = op(e4, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP23])])).
fof(f1479, plain, (sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1478, f977])).
fof(f977, plain, ~ sP24, inference(subsumption_resolution, [], [f453, f193])).
fof(f193, plain, (e2 = op(e4, e3)), inference(cnf_transformation, [], [f2])).
fof(f453, plain, (~ (e2 = op(e4, e3)) | ~ sP24), inference(cnf_transformation, [], [f135])).
fof(f135, plain, ((~ (e4 = op(e4, e3)) & ~ (e3 = op(e4, e3)) & ~ (e2 = op(e4, e3)) & ~ (e1 = op(e4, e3)) & ~ (e0 = op(e4, e3))) | ~ sP24), inference(nnf_transformation, [], [f31])).
fof(f31, plain, ((~ (e4 = op(e4, e3)) & ~ (e3 = op(e4, e3)) & ~ (e2 = op(e4, e3)) & ~ (e1 = op(e4, e3)) & ~ (e0 = op(e4, e3))) | ~ sP24), inference(usedef, [], [e31])).
fof(e31, plain, (sP24 <=> (~ (e4 = op(e4, e3)) & ~ (e3 = op(e4, e3)) & ~ (e2 = op(e4, e3)) & ~ (e1 = op(e4, e3)) & ~ (e0 = op(e4, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP24])])).
fof(f1478, plain, (sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1477, f974])).
fof(f974, plain, ~ sP25, inference(subsumption_resolution, [], [f447, f194])).
fof(f194, plain, (e1 = op(e4, e4)), inference(cnf_transformation, [], [f2])).
fof(f447, plain, (~ (e1 = op(e4, e4)) | ~ sP25), inference(cnf_transformation, [], [f134])).
fof(f134, plain, ((~ (e4 = op(e4, e4)) & ~ (e3 = op(e4, e4)) & ~ (e2 = op(e4, e4)) & ~ (e1 = op(e4, e4)) & ~ (e0 = op(e4, e4))) | ~ sP25), inference(nnf_transformation, [], [f32])).
fof(f32, plain, ((~ (e4 = op(e4, e4)) & ~ (e3 = op(e4, e4)) & ~ (e2 = op(e4, e4)) & ~ (e1 = op(e4, e4)) & ~ (e0 = op(e4, e4))) | ~ sP25), inference(usedef, [], [e32])).
fof(e32, plain, (sP25 <=> (~ (e4 = op(e4, e4)) & ~ (e3 = op(e4, e4)) & ~ (e2 = op(e4, e4)) & ~ (e1 = op(e4, e4)) & ~ (e0 = op(e4, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP25])])).
fof(f1477, plain, (sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1476, f617])).
fof(f1476, plain, (~ (unit = op(unit, unit)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1475, f616])).
fof(f1475, plain, (~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1474, f612])).
fof(f1474, plain, (~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1473, f615])).
fof(f1473, plain, (~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1472, f610])).
fof(f1472, plain, (~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1471, f614])).
fof(f1471, plain, (~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1470, f608])).
fof(f1470, plain, (~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1469, f606])).
fof(f1469, plain, (~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1468, f825])).
fof(f825, plain, ~ sP26, inference(trivial_inequality_removal, [], [f739])).
fof(f739, plain, (~ (unit = unit) | ~ sP26), inference(definition_unfolding, [], [f441, f195])).
fof(f441, plain, (~ (e0 = unit) | ~ sP26), inference(cnf_transformation, [], [f133])).
fof(f133, plain, ((~ (e4 = unit) & ~ (e3 = unit) & ~ (e2 = unit) & ~ (e1 = unit) & ~ (e0 = unit)) | ~ sP26), inference(nnf_transformation, [], [f33])).
fof(f33, plain, ((~ (e4 = unit) & ~ (e3 = unit) & ~ (e2 = unit) & ~ (e1 = unit) & ~ (e0 = unit)) | ~ sP26), inference(usedef, [], [e33])).
fof(e33, plain, (sP26 <=> (~ (e4 = unit) & ~ (e3 = unit) & ~ (e2 = unit) & ~ (e1 = unit) & ~ (e0 = unit))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP26])])).
fof(f1468, plain, (sP26 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1467, f970])).
fof(f970, plain, ~ sP27, inference(subsumption_resolution, [], [f738, f617])).
fof(f738, plain, (~ (unit = op(unit, unit)) | ~ sP27), inference(definition_unfolding, [], [f436, f195, f195, f195])).
fof(f436, plain, (~ (e0 = op(e0, e0)) | ~ sP27), inference(cnf_transformation, [], [f132])).
fof(f132, plain, ((~ (e0 = op(e0, e4)) & ~ (e0 = op(e0, e3)) & ~ (e0 = op(e0, e2)) & ~ (e0 = op(e0, e1)) & ~ (e0 = op(e0, e0))) | ~ sP27), inference(nnf_transformation, [], [f34])).
fof(f34, plain, ((~ (e0 = op(e0, e4)) & ~ (e0 = op(e0, e3)) & ~ (e0 = op(e0, e2)) & ~ (e0 = op(e0, e1)) & ~ (e0 = op(e0, e0))) | ~ sP27), inference(usedef, [], [e34])).
fof(e34, plain, (sP27 <=> (~ (e0 = op(e0, e4)) & ~ (e0 = op(e0, e3)) & ~ (e0 = op(e0, e2)) & ~ (e0 = op(e0, e1)) & ~ (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP27])])).
fof(f1467, plain, (sP27 | sP26 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1466, f965])).
fof(f965, plain, ~ sP28, inference(subsumption_resolution, [], [f733, f617])).
fof(f733, plain, (~ (unit = op(unit, unit)) | ~ sP28), inference(definition_unfolding, [], [f431, f195, f195, f195])).
fof(f431, plain, (~ (e0 = op(e0, e0)) | ~ sP28), inference(cnf_transformation, [], [f131])).
fof(f131, plain, ((~ (e0 = op(e4, e0)) & ~ (e0 = op(e3, e0)) & ~ (e0 = op(e2, e0)) & ~ (e0 = op(e1, e0)) & ~ (e0 = op(e0, e0))) | ~ sP28), inference(nnf_transformation, [], [f35])).
fof(f35, plain, ((~ (e0 = op(e4, e0)) & ~ (e0 = op(e3, e0)) & ~ (e0 = op(e2, e0)) & ~ (e0 = op(e1, e0)) & ~ (e0 = op(e0, e0))) | ~ sP28), inference(usedef, [], [e35])).
fof(e35, plain, (sP28 <=> (~ (e0 = op(e4, e0)) & ~ (e0 = op(e3, e0)) & ~ (e0 = op(e2, e0)) & ~ (e0 = op(e1, e0)) & ~ (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP28])])).
fof(f1466, plain, (sP28 | sP27 | sP26 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1465, f960])).
fof(f960, plain, ~ sP29, inference(subsumption_resolution, [], [f727, f616])).
fof(f727, plain, (~ (e1 = op(unit, e1)) | ~ sP29), inference(definition_unfolding, [], [f427, f195])).
fof(f427, plain, (~ (e1 = op(e0, e1)) | ~ sP29), inference(cnf_transformation, [], [f130])).
fof(f130, plain, ((~ (e1 = op(e0, e4)) & ~ (e1 = op(e0, e3)) & ~ (e1 = op(e0, e2)) & ~ (e1 = op(e0, e1)) & ~ (e1 = op(e0, e0))) | ~ sP29), inference(nnf_transformation, [], [f36])).
fof(f36, plain, ((~ (e1 = op(e0, e4)) & ~ (e1 = op(e0, e3)) & ~ (e1 = op(e0, e2)) & ~ (e1 = op(e0, e1)) & ~ (e1 = op(e0, e0))) | ~ sP29), inference(usedef, [], [e36])).
fof(e36, plain, (sP29 <=> (~ (e1 = op(e0, e4)) & ~ (e1 = op(e0, e3)) & ~ (e1 = op(e0, e2)) & ~ (e1 = op(e0, e1)) & ~ (e1 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP29])])).
fof(f1465, plain, (sP29 | sP28 | sP27 | sP26 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1464, f956])).
fof(f956, plain, ~ sP30, inference(subsumption_resolution, [], [f722, f612])).
fof(f722, plain, (~ (e1 = op(e1, unit)) | ~ sP30), inference(definition_unfolding, [], [f422, f195])).
fof(f422, plain, (~ (e1 = op(e1, e0)) | ~ sP30), inference(cnf_transformation, [], [f129])).
fof(f129, plain, ((~ (e1 = op(e4, e0)) & ~ (e1 = op(e3, e0)) & ~ (e1 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e1 = op(e0, e0))) | ~ sP30), inference(nnf_transformation, [], [f37])).
fof(f37, plain, ((~ (e1 = op(e4, e0)) & ~ (e1 = op(e3, e0)) & ~ (e1 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e1 = op(e0, e0))) | ~ sP30), inference(usedef, [], [e37])).
fof(e37, plain, (sP30 <=> (~ (e1 = op(e4, e0)) & ~ (e1 = op(e3, e0)) & ~ (e1 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e1 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP30])])).
fof(f1464, plain, (sP30 | sP29 | sP28 | sP27 | sP26 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1463, f952])).
fof(f952, plain, ~ sP31, inference(subsumption_resolution, [], [f716, f615])).
fof(f716, plain, (~ (e2 = op(unit, e2)) | ~ sP31), inference(definition_unfolding, [], [f418, f195])).
fof(f418, plain, (~ (e2 = op(e0, e2)) | ~ sP31), inference(cnf_transformation, [], [f128])).
fof(f128, plain, ((~ (e2 = op(e0, e4)) & ~ (e2 = op(e0, e3)) & ~ (e2 = op(e0, e2)) & ~ (e2 = op(e0, e1)) & ~ (e2 = op(e0, e0))) | ~ sP31), inference(nnf_transformation, [], [f38])).
fof(f38, plain, ((~ (e2 = op(e0, e4)) & ~ (e2 = op(e0, e3)) & ~ (e2 = op(e0, e2)) & ~ (e2 = op(e0, e1)) & ~ (e2 = op(e0, e0))) | ~ sP31), inference(usedef, [], [e38])).
fof(e38, plain, (sP31 <=> (~ (e2 = op(e0, e4)) & ~ (e2 = op(e0, e3)) & ~ (e2 = op(e0, e2)) & ~ (e2 = op(e0, e1)) & ~ (e2 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP31])])).
fof(f1463, plain, (sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1462, f949])).
fof(f949, plain, ~ sP32, inference(subsumption_resolution, [], [f711, f610])).
fof(f711, plain, (~ (e2 = op(e2, unit)) | ~ sP32), inference(definition_unfolding, [], [f413, f195])).
fof(f413, plain, (~ (e2 = op(e2, e0)) | ~ sP32), inference(cnf_transformation, [], [f127])).
fof(f127, plain, ((~ (e2 = op(e4, e0)) & ~ (e2 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e2 = op(e1, e0)) & ~ (e2 = op(e0, e0))) | ~ sP32), inference(nnf_transformation, [], [f39])).
fof(f39, plain, ((~ (e2 = op(e4, e0)) & ~ (e2 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e2 = op(e1, e0)) & ~ (e2 = op(e0, e0))) | ~ sP32), inference(usedef, [], [e39])).
fof(e39, plain, (sP32 <=> (~ (e2 = op(e4, e0)) & ~ (e2 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e2 = op(e1, e0)) & ~ (e2 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP32])])).
fof(f1462, plain, (sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1461, f946])).
fof(f946, plain, ~ sP33, inference(subsumption_resolution, [], [f705, f614])).
fof(f705, plain, (~ (e3 = op(unit, e3)) | ~ sP33), inference(definition_unfolding, [], [f409, f195])).
fof(f409, plain, (~ (e3 = op(e0, e3)) | ~ sP33), inference(cnf_transformation, [], [f126])).
fof(f126, plain, ((~ (e3 = op(e0, e4)) & ~ (e3 = op(e0, e3)) & ~ (e3 = op(e0, e2)) & ~ (e3 = op(e0, e1)) & ~ (e3 = op(e0, e0))) | ~ sP33), inference(nnf_transformation, [], [f40])).
fof(f40, plain, ((~ (e3 = op(e0, e4)) & ~ (e3 = op(e0, e3)) & ~ (e3 = op(e0, e2)) & ~ (e3 = op(e0, e1)) & ~ (e3 = op(e0, e0))) | ~ sP33), inference(usedef, [], [e40])).
fof(e40, plain, (sP33 <=> (~ (e3 = op(e0, e4)) & ~ (e3 = op(e0, e3)) & ~ (e3 = op(e0, e2)) & ~ (e3 = op(e0, e1)) & ~ (e3 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP33])])).
fof(f1461, plain, (sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1460, f944])).
fof(f944, plain, ~ sP34, inference(subsumption_resolution, [], [f700, f608])).
fof(f700, plain, (~ (e3 = op(e3, unit)) | ~ sP34), inference(definition_unfolding, [], [f404, f195])).
fof(f404, plain, (~ (e3 = op(e3, e0)) | ~ sP34), inference(cnf_transformation, [], [f125])).
fof(f125, plain, ((~ (e3 = op(e4, e0)) & ~ (e3 = op(e3, e0)) & ~ (e3 = op(e2, e0)) & ~ (e3 = op(e1, e0)) & ~ (e3 = op(e0, e0))) | ~ sP34), inference(nnf_transformation, [], [f41])).
fof(f41, plain, ((~ (e3 = op(e4, e0)) & ~ (e3 = op(e3, e0)) & ~ (e3 = op(e2, e0)) & ~ (e3 = op(e1, e0)) & ~ (e3 = op(e0, e0))) | ~ sP34), inference(usedef, [], [e41])).
fof(e41, plain, (sP34 <=> (~ (e3 = op(e4, e0)) & ~ (e3 = op(e3, e0)) & ~ (e3 = op(e2, e0)) & ~ (e3 = op(e1, e0)) & ~ (e3 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP34])])).
fof(f1460, plain, (sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1459, f942])).
fof(f942, plain, ~ sP35, inference(subsumption_resolution, [], [f694, f613])).
fof(f694, plain, (~ (e4 = op(unit, e4)) | ~ sP35), inference(definition_unfolding, [], [f400, f195])).
fof(f400, plain, (~ (e4 = op(e0, e4)) | ~ sP35), inference(cnf_transformation, [], [f124])).
fof(f124, plain, ((~ (e4 = op(e0, e4)) & ~ (e4 = op(e0, e3)) & ~ (e4 = op(e0, e2)) & ~ (e4 = op(e0, e1)) & ~ (e4 = op(e0, e0))) | ~ sP35), inference(nnf_transformation, [], [f42])).
fof(f42, plain, ((~ (e4 = op(e0, e4)) & ~ (e4 = op(e0, e3)) & ~ (e4 = op(e0, e2)) & ~ (e4 = op(e0, e1)) & ~ (e4 = op(e0, e0))) | ~ sP35), inference(usedef, [], [e42])).
fof(e42, plain, (sP35 <=> (~ (e4 = op(e0, e4)) & ~ (e4 = op(e0, e3)) & ~ (e4 = op(e0, e2)) & ~ (e4 = op(e0, e1)) & ~ (e4 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP35])])).
fof(f1459, plain, (sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1458, f941])).
fof(f941, plain, ~ sP36, inference(subsumption_resolution, [], [f689, f606])).
fof(f689, plain, (~ (e4 = op(e4, unit)) | ~ sP36), inference(definition_unfolding, [], [f395, f195])).
fof(f395, plain, (~ (e4 = op(e4, e0)) | ~ sP36), inference(cnf_transformation, [], [f123])).
fof(f123, plain, ((~ (e4 = op(e4, e0)) & ~ (e4 = op(e3, e0)) & ~ (e4 = op(e2, e0)) & ~ (e4 = op(e1, e0)) & ~ (e4 = op(e0, e0))) | ~ sP36), inference(nnf_transformation, [], [f43])).
fof(f43, plain, ((~ (e4 = op(e4, e0)) & ~ (e4 = op(e3, e0)) & ~ (e4 = op(e2, e0)) & ~ (e4 = op(e1, e0)) & ~ (e4 = op(e0, e0))) | ~ sP36), inference(usedef, [], [e43])).
fof(e43, plain, (sP36 <=> (~ (e4 = op(e4, e0)) & ~ (e4 = op(e3, e0)) & ~ (e4 = op(e2, e0)) & ~ (e4 = op(e1, e0)) & ~ (e4 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP36])])).
fof(f1458, plain, (sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1457, f940])).
fof(f940, plain, ~ sP37, inference(subsumption_resolution, [], [f684, f611])).
fof(f684, plain, (~ (op(e1, e4) = unit) | ~ sP37), inference(definition_unfolding, [], [f390, f195])).
fof(f390, plain, (~ (e0 = op(e1, e4)) | ~ sP37), inference(cnf_transformation, [], [f122])).
fof(f122, plain, ((~ (e0 = op(e1, e4)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e1, e0))) | ~ sP37), inference(nnf_transformation, [], [f44])).
fof(f44, plain, ((~ (e0 = op(e1, e4)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e1, e0))) | ~ sP37), inference(usedef, [], [e44])).
fof(e44, plain, (sP37 <=> (~ (e0 = op(e1, e4)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP37])])).
fof(f1457, plain, (sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1456, f939])).
fof(f939, plain, ~ sP38, inference(subsumption_resolution, [], [f680, f607])).
fof(f680, plain, (~ (op(e3, e1) = unit) | ~ sP38), inference(definition_unfolding, [], [f384, f195])).
fof(f384, plain, (~ (e0 = op(e3, e1)) | ~ sP38), inference(cnf_transformation, [], [f121])).
fof(f121, plain, ((~ (e0 = op(e4, e1)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e0, e1))) | ~ sP38), inference(nnf_transformation, [], [f45])).
fof(f45, plain, ((~ (e0 = op(e4, e1)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e0, e1))) | ~ sP38), inference(usedef, [], [e45])).
fof(e45, plain, (sP38 <=> (~ (e0 = op(e4, e1)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP38])])).
fof(f1456, plain, (sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1455, f937])).
fof(f937, plain, ~ sP39, inference(subsumption_resolution, [], [f678, f612])).
fof(f678, plain, (~ (e1 = op(e1, unit)) | ~ sP39), inference(definition_unfolding, [], [f376, f195])).
fof(f376, plain, (~ (e1 = op(e1, e0)) | ~ sP39), inference(cnf_transformation, [], [f120])).
fof(f120, plain, ((~ (e1 = op(e1, e4)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e1, e0))) | ~ sP39), inference(nnf_transformation, [], [f46])).
fof(f46, plain, ((~ (e1 = op(e1, e4)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e1, e0))) | ~ sP39), inference(usedef, [], [e46])).
fof(e46, plain, (sP39 <=> (~ (e1 = op(e1, e4)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP39])])).
fof(f1455, plain, (sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1454, f932])).
fof(f932, plain, ~ sP40, inference(subsumption_resolution, [], [f677, f616])).
fof(f677, plain, (~ (e1 = op(unit, e1)) | ~ sP40), inference(definition_unfolding, [], [f371, f195])).
fof(f371, plain, (~ (e1 = op(e0, e1)) | ~ sP40), inference(cnf_transformation, [], [f119])).
fof(f119, plain, ((~ (e1 = op(e4, e1)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e0, e1))) | ~ sP40), inference(nnf_transformation, [], [f47])).
fof(f47, plain, ((~ (e1 = op(e4, e1)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e0, e1))) | ~ sP40), inference(usedef, [], [e47])).
fof(e47, plain, (sP40 <=> (~ (e1 = op(e4, e1)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP40])])).
fof(f1454, plain, (sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1453, f927])).
fof(f927, plain, ~ sP41, inference(subsumption_resolution, [], [f367, f176])).
fof(f367, plain, (~ (e2 = op(e1, e1)) | ~ sP41), inference(cnf_transformation, [], [f118])).
fof(f118, plain, ((~ (e2 = op(e1, e4)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e1, e0))) | ~ sP41), inference(nnf_transformation, [], [f48])).
fof(f48, plain, ((~ (e2 = op(e1, e4)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e1, e0))) | ~ sP41), inference(usedef, [], [e48])).
fof(e48, plain, (sP41 <=> (~ (e2 = op(e1, e4)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP41])])).
fof(f1453, plain, (sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1452, f923])).
fof(f923, plain, ~ sP42, inference(subsumption_resolution, [], [f362, f176])).
fof(f362, plain, (~ (e2 = op(e1, e1)) | ~ sP42), inference(cnf_transformation, [], [f117])).
fof(f117, plain, ((~ (e2 = op(e4, e1)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e0, e1))) | ~ sP42), inference(nnf_transformation, [], [f49])).
fof(f49, plain, ((~ (e2 = op(e4, e1)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e0, e1))) | ~ sP42), inference(usedef, [], [e49])).
fof(e49, plain, (sP42 <=> (~ (e2 = op(e4, e1)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP42])])).
fof(f1452, plain, (sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1451, f919])).
fof(f919, plain, ~ sP43, inference(subsumption_resolution, [], [f358, f177])).
fof(f358, plain, (~ (e3 = op(e1, e2)) | ~ sP43), inference(cnf_transformation, [], [f116])).
fof(f116, plain, ((~ (e3 = op(e1, e4)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e1, e0))) | ~ sP43), inference(nnf_transformation, [], [f50])).
fof(f50, plain, ((~ (e3 = op(e1, e4)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e1, e0))) | ~ sP43), inference(usedef, [], [e50])).
fof(e50, plain, (sP43 <=> (~ (e3 = op(e1, e4)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP43])])).
fof(f1451, plain, (sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1450, f916])).
fof(f916, plain, ~ sP44, inference(subsumption_resolution, [], [f355, f191])).
fof(f355, plain, (~ (e3 = op(e4, e1)) | ~ sP44), inference(cnf_transformation, [], [f115])).
fof(f115, plain, ((~ (e3 = op(e4, e1)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e0, e1))) | ~ sP44), inference(nnf_transformation, [], [f51])).
fof(f51, plain, ((~ (e3 = op(e4, e1)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e0, e1))) | ~ sP44), inference(usedef, [], [e51])).
fof(e51, plain, (sP44 <=> (~ (e3 = op(e4, e1)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP44])])).
fof(f1450, plain, (sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1449, f915])).
fof(f915, plain, ~ sP45, inference(subsumption_resolution, [], [f349, f178])).
fof(f349, plain, (~ (e4 = op(e1, e3)) | ~ sP45), inference(cnf_transformation, [], [f114])).
fof(f114, plain, ((~ (e4 = op(e1, e4)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e1, e0))) | ~ sP45), inference(nnf_transformation, [], [f52])).
fof(f52, plain, ((~ (e4 = op(e1, e4)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e1, e0))) | ~ sP45), inference(usedef, [], [e52])).
fof(e52, plain, (sP45 <=> (~ (e4 = op(e1, e4)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP45])])).
fof(f1449, plain, (sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1448, f913])).
fof(f913, plain, ~ sP46, inference(subsumption_resolution, [], [f343, f181])).
fof(f343, plain, (~ (e4 = op(e2, e1)) | ~ sP46), inference(cnf_transformation, [], [f113])).
fof(f113, plain, ((~ (e4 = op(e4, e1)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e0, e1))) | ~ sP46), inference(nnf_transformation, [], [f53])).
fof(f53, plain, ((~ (e4 = op(e4, e1)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e0, e1))) | ~ sP46), inference(usedef, [], [e53])).
fof(e53, plain, (sP46 <=> (~ (e4 = op(e4, e1)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP46])])).
fof(f1448, plain, (sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1447, f910])).
fof(f910, plain, ~ sP47, inference(subsumption_resolution, [], [f667, f609])).
fof(f667, plain, (~ (op(e2, e3) = unit) | ~ sP47), inference(definition_unfolding, [], [f339, f195])).
fof(f339, plain, (~ (e0 = op(e2, e3)) | ~ sP47), inference(cnf_transformation, [], [f112])).
fof(f112, plain, ((~ (e0 = op(e2, e4)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e2, e0))) | ~ sP47), inference(nnf_transformation, [], [f54])).
fof(f54, plain, ((~ (e0 = op(e2, e4)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e2, e0))) | ~ sP47), inference(usedef, [], [e54])).
fof(e54, plain, (sP47 <=> (~ (e0 = op(e2, e4)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP47])])).
fof(f1447, plain, (sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1446, f908])).
fof(f908, plain, ~ sP48, inference(subsumption_resolution, [], [f661, f605])).
fof(f661, plain, (~ (op(e4, e2) = unit) | ~ sP48), inference(definition_unfolding, [], [f335, f195])).
fof(f335, plain, (~ (e0 = op(e4, e2)) | ~ sP48), inference(cnf_transformation, [], [f111])).
fof(f111, plain, ((~ (e0 = op(e4, e2)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e0, e2))) | ~ sP48), inference(nnf_transformation, [], [f55])).
fof(f55, plain, ((~ (e0 = op(e4, e2)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e0, e2))) | ~ sP48), inference(usedef, [], [e55])).
fof(e55, plain, (sP48 <=> (~ (e0 = op(e4, e2)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP48])])).
fof(f1446, plain, (sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1445, f907])).
fof(f907, plain, ~ sP49, inference(subsumption_resolution, [], [f328, f182])).
fof(f328, plain, (~ (e1 = op(e2, e2)) | ~ sP49), inference(cnf_transformation, [], [f110])).
fof(f110, plain, ((~ (e1 = op(e2, e4)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e2, e0))) | ~ sP49), inference(nnf_transformation, [], [f56])).
fof(f56, plain, ((~ (e1 = op(e2, e4)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e2, e0))) | ~ sP49), inference(usedef, [], [e56])).
fof(e56, plain, (sP49 <=> (~ (e1 = op(e2, e4)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP49])])).
fof(f1445, plain, (sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1444, f904])).
fof(f904, plain, ~ sP50, inference(subsumption_resolution, [], [f323, f182])).
fof(f323, plain, (~ (e1 = op(e2, e2)) | ~ sP50), inference(cnf_transformation, [], [f109])).
fof(f109, plain, ((~ (e1 = op(e4, e2)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e0, e2))) | ~ sP50), inference(nnf_transformation, [], [f57])).
fof(f57, plain, ((~ (e1 = op(e4, e2)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e0, e2))) | ~ sP50), inference(usedef, [], [e57])).
fof(e57, plain, (sP50 <=> (~ (e1 = op(e4, e2)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP50])])).
fof(f1444, plain, (sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1443, f901])).
fof(f901, plain, ~ sP51, inference(subsumption_resolution, [], [f658, f610])).
fof(f658, plain, (~ (e2 = op(e2, unit)) | ~ sP51), inference(definition_unfolding, [], [f316, f195])).
fof(f316, plain, (~ (e2 = op(e2, e0)) | ~ sP51), inference(cnf_transformation, [], [f108])).
fof(f108, plain, ((~ (e2 = op(e2, e4)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e2, e0))) | ~ sP51), inference(nnf_transformation, [], [f58])).
fof(f58, plain, ((~ (e2 = op(e2, e4)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e2, e0))) | ~ sP51), inference(usedef, [], [e58])).
fof(e58, plain, (sP51 <=> (~ (e2 = op(e2, e4)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP51])])).
fof(f1443, plain, (sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1442, f896])).
fof(f896, plain, ~ sP52, inference(subsumption_resolution, [], [f657, f615])).
fof(f657, plain, (~ (e2 = op(unit, e2)) | ~ sP52), inference(definition_unfolding, [], [f311, f195])).
fof(f311, plain, (~ (e2 = op(e0, e2)) | ~ sP52), inference(cnf_transformation, [], [f107])).
fof(f107, plain, ((~ (e2 = op(e4, e2)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e0, e2))) | ~ sP52), inference(nnf_transformation, [], [f59])).
fof(f59, plain, ((~ (e2 = op(e4, e2)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e0, e2))) | ~ sP52), inference(usedef, [], [e59])).
fof(e59, plain, (sP52 <=> (~ (e2 = op(e4, e2)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP52])])).
fof(f1442, plain, (sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1441, f891])).
fof(f891, plain, ~ sP53, inference(subsumption_resolution, [], [f310, f184])).
fof(f310, plain, (~ (e3 = op(e2, e4)) | ~ sP53), inference(cnf_transformation, [], [f106])).
fof(f106, plain, ((~ (e3 = op(e2, e4)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e2, e0))) | ~ sP53), inference(nnf_transformation, [], [f60])).
fof(f60, plain, ((~ (e3 = op(e2, e4)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e2, e0))) | ~ sP53), inference(usedef, [], [e60])).
fof(e60, plain, (sP53 <=> (~ (e3 = op(e2, e4)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP53])])).
fof(f1441, plain, (sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1440, f890])).
fof(f890, plain, ~ sP54, inference(subsumption_resolution, [], [f302, f177])).
fof(f302, plain, (~ (e3 = op(e1, e2)) | ~ sP54), inference(cnf_transformation, [], [f105])).
fof(f105, plain, ((~ (e3 = op(e4, e2)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e0, e2))) | ~ sP54), inference(nnf_transformation, [], [f61])).
fof(f61, plain, ((~ (e3 = op(e4, e2)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e0, e2))) | ~ sP54), inference(usedef, [], [e61])).
fof(e61, plain, (sP54 <=> (~ (e3 = op(e4, e2)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP54])])).
fof(f1440, plain, (sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1439, f886])).
fof(f886, plain, ~ sP55, inference(subsumption_resolution, [], [f297, f181])).
fof(f297, plain, (~ (e4 = op(e2, e1)) | ~ sP55), inference(cnf_transformation, [], [f104])).
fof(f104, plain, ((~ (e4 = op(e2, e4)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e2, e0))) | ~ sP55), inference(nnf_transformation, [], [f62])).
fof(f62, plain, ((~ (e4 = op(e2, e4)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e2, e0))) | ~ sP55), inference(usedef, [], [e62])).
fof(e62, plain, (sP55 <=> (~ (e4 = op(e2, e4)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP55])])).
fof(f1439, plain, (sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1438, f882])).
fof(f882, plain, ~ sP56, inference(subsumption_resolution, [], [f294, f187])).
fof(f294, plain, (~ (e4 = op(e3, e2)) | ~ sP56), inference(cnf_transformation, [], [f103])).
fof(f103, plain, ((~ (e4 = op(e4, e2)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e0, e2))) | ~ sP56), inference(nnf_transformation, [], [f63])).
fof(f63, plain, ((~ (e4 = op(e4, e2)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e0, e2))) | ~ sP56), inference(usedef, [], [e63])).
fof(e63, plain, (sP56 <=> (~ (e4 = op(e4, e2)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP56])])).
fof(f1438, plain, (sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1437, f880])).
fof(f880, plain, ~ sP57, inference(subsumption_resolution, [], [f651, f607])).
fof(f651, plain, (~ (op(e3, e1) = unit) | ~ sP57), inference(definition_unfolding, [], [f287, f195])).
fof(f287, plain, (~ (e0 = op(e3, e1)) | ~ sP57), inference(cnf_transformation, [], [f102])).
fof(f102, plain, ((~ (e0 = op(e3, e4)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e3, e0))) | ~ sP57), inference(nnf_transformation, [], [f64])).
fof(f64, plain, ((~ (e0 = op(e3, e4)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e3, e0))) | ~ sP57), inference(usedef, [], [e64])).
fof(e64, plain, (sP57 <=> (~ (e0 = op(e3, e4)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP57])])).
fof(f1437, plain, (sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1436, f876])).
fof(f876, plain, ~ sP58, inference(subsumption_resolution, [], [f645, f609])).
fof(f645, plain, (~ (op(e2, e3) = unit) | ~ sP58), inference(definition_unfolding, [], [f283, f195])).
fof(f283, plain, (~ (e0 = op(e2, e3)) | ~ sP58), inference(cnf_transformation, [], [f101])).
fof(f101, plain, ((~ (e0 = op(e4, e3)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e0, e3))) | ~ sP58), inference(nnf_transformation, [], [f65])).
fof(f65, plain, ((~ (e0 = op(e4, e3)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e0, e3))) | ~ sP58), inference(usedef, [], [e65])).
fof(e65, plain, (sP58 <=> (~ (e0 = op(e4, e3)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP58])])).
fof(f1436, plain, (sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1435, f873])).
fof(f873, plain, ~ sP59, inference(subsumption_resolution, [], [f279, f188])).
fof(f279, plain, (~ (e1 = op(e3, e3)) | ~ sP59), inference(cnf_transformation, [], [f100])).
fof(f100, plain, ((~ (e1 = op(e3, e4)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e3, e0))) | ~ sP59), inference(nnf_transformation, [], [f66])).
fof(f66, plain, ((~ (e1 = op(e3, e4)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e3, e0))) | ~ sP59), inference(usedef, [], [e66])).
fof(e66, plain, (sP59 <=> (~ (e1 = op(e3, e4)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP59])])).
fof(f1435, plain, (sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1434, f871])).
fof(f871, plain, ~ sP60, inference(subsumption_resolution, [], [f274, f188])).
fof(f274, plain, (~ (e1 = op(e3, e3)) | ~ sP60), inference(cnf_transformation, [], [f99])).
fof(f99, plain, ((~ (e1 = op(e4, e3)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e0, e3))) | ~ sP60), inference(nnf_transformation, [], [f67])).
fof(f67, plain, ((~ (e1 = op(e4, e3)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e0, e3))) | ~ sP60), inference(usedef, [], [e67])).
fof(e67, plain, (sP60 <=> (~ (e1 = op(e4, e3)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP60])])).
fof(f1434, plain, (sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1433, f869])).
fof(f869, plain, ~ sP61, inference(subsumption_resolution, [], [f270, f189])).
fof(f270, plain, (~ (e2 = op(e3, e4)) | ~ sP61), inference(cnf_transformation, [], [f98])).
fof(f98, plain, ((~ (e2 = op(e3, e4)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e3, e0))) | ~ sP61), inference(nnf_transformation, [], [f68])).
fof(f68, plain, ((~ (e2 = op(e3, e4)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e3, e0))) | ~ sP61), inference(usedef, [], [e68])).
fof(e68, plain, (sP61 <=> (~ (e2 = op(e3, e4)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP61])])).
fof(f1433, plain, (sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1432, f868])).
fof(f868, plain, ~ sP62, inference(subsumption_resolution, [], [f265, f193])).
fof(f265, plain, (~ (e2 = op(e4, e3)) | ~ sP62), inference(cnf_transformation, [], [f97])).
fof(f97, plain, ((~ (e2 = op(e4, e3)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e0, e3))) | ~ sP62), inference(nnf_transformation, [], [f69])).
fof(f69, plain, ((~ (e2 = op(e4, e3)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e0, e3))) | ~ sP62), inference(usedef, [], [e69])).
fof(e69, plain, (sP62 <=> (~ (e2 = op(e4, e3)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP62])])).
fof(f1432, plain, (sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1431, f867])).
fof(f867, plain, ~ sP63, inference(subsumption_resolution, [], [f638, f608])).
fof(f638, plain, (~ (e3 = op(e3, unit)) | ~ sP63), inference(definition_unfolding, [], [f256, f195])).
fof(f256, plain, (~ (e3 = op(e3, e0)) | ~ sP63), inference(cnf_transformation, [], [f96])).
fof(f96, plain, ((~ (e3 = op(e3, e4)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e3, e0))) | ~ sP63), inference(nnf_transformation, [], [f70])).
fof(f70, plain, ((~ (e3 = op(e3, e4)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e3, e0))) | ~ sP63), inference(usedef, [], [e70])).
fof(e70, plain, (sP63 <=> (~ (e3 = op(e3, e4)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP63])])).
fof(f1431, plain, (sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1430, f862])).
fof(f862, plain, ~ sP64, inference(subsumption_resolution, [], [f637, f614])).
fof(f637, plain, (~ (e3 = op(unit, e3)) | ~ sP64), inference(definition_unfolding, [], [f251, f195])).
fof(f251, plain, (~ (e3 = op(e0, e3)) | ~ sP64), inference(cnf_transformation, [], [f95])).
fof(f95, plain, ((~ (e3 = op(e4, e3)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e0, e3))) | ~ sP64), inference(nnf_transformation, [], [f71])).
fof(f71, plain, ((~ (e3 = op(e4, e3)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e0, e3))) | ~ sP64), inference(usedef, [], [e71])).
fof(e71, plain, (sP64 <=> (~ (e3 = op(e4, e3)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP64])])).
fof(f1430, plain, (sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1429, f857])).
fof(f857, plain, ~ sP65, inference(subsumption_resolution, [], [f248, f187])).
fof(f248, plain, (~ (e4 = op(e3, e2)) | ~ sP65), inference(cnf_transformation, [], [f94])).
fof(f94, plain, ((~ (e4 = op(e3, e4)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e3, e0))) | ~ sP65), inference(nnf_transformation, [], [f72])).
fof(f72, plain, ((~ (e4 = op(e3, e4)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e3, e0))) | ~ sP65), inference(usedef, [], [e72])).
fof(e72, plain, (sP65 <=> (~ (e4 = op(e3, e4)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP65])])).
fof(f1429, plain, (sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1428, f854])).
fof(f854, plain, ~ sP66, inference(subsumption_resolution, [], [f242, f178])).
fof(f242, plain, (~ (e4 = op(e1, e3)) | ~ sP66), inference(cnf_transformation, [], [f93])).
fof(f93, plain, ((~ (e4 = op(e4, e3)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e0, e3))) | ~ sP66), inference(nnf_transformation, [], [f73])).
fof(f73, plain, ((~ (e4 = op(e4, e3)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e0, e3))) | ~ sP66), inference(usedef, [], [e73])).
fof(e73, plain, (sP66 <=> (~ (e4 = op(e4, e3)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP66])])).
fof(f1428, plain, (sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1427, f850])).
fof(f850, plain, ~ sP67, inference(subsumption_resolution, [], [f632, f605])).
fof(f632, plain, (~ (op(e4, e2) = unit) | ~ sP67), inference(definition_unfolding, [], [f238, f195])).
fof(f238, plain, (~ (e0 = op(e4, e2)) | ~ sP67), inference(cnf_transformation, [], [f92])).
fof(f92, plain, ((~ (e0 = op(e4, e4)) & ~ (e0 = op(e4, e3)) & ~ (e0 = op(e4, e2)) & ~ (e0 = op(e4, e1)) & ~ (e0 = op(e4, e0))) | ~ sP67), inference(nnf_transformation, [], [f74])).
fof(f74, plain, ((~ (e0 = op(e4, e4)) & ~ (e0 = op(e4, e3)) & ~ (e0 = op(e4, e2)) & ~ (e0 = op(e4, e1)) & ~ (e0 = op(e4, e0))) | ~ sP67), inference(usedef, [], [e74])).
fof(e74, plain, (sP67 <=> (~ (e0 = op(e4, e4)) & ~ (e0 = op(e4, e3)) & ~ (e0 = op(e4, e2)) & ~ (e0 = op(e4, e1)) & ~ (e0 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP67])])).
fof(f1427, plain, (sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1426, f847])).
fof(f847, plain, ~ sP68, inference(subsumption_resolution, [], [f628, f611])).
fof(f628, plain, (~ (op(e1, e4) = unit) | ~ sP68), inference(definition_unfolding, [], [f232, f195])).
fof(f232, plain, (~ (e0 = op(e1, e4)) | ~ sP68), inference(cnf_transformation, [], [f91])).
fof(f91, plain, ((~ (e0 = op(e4, e4)) & ~ (e0 = op(e3, e4)) & ~ (e0 = op(e2, e4)) & ~ (e0 = op(e1, e4)) & ~ (e0 = op(e0, e4))) | ~ sP68), inference(nnf_transformation, [], [f75])).
fof(f75, plain, ((~ (e0 = op(e4, e4)) & ~ (e0 = op(e3, e4)) & ~ (e0 = op(e2, e4)) & ~ (e0 = op(e1, e4)) & ~ (e0 = op(e0, e4))) | ~ sP68), inference(usedef, [], [e75])).
fof(e75, plain, (sP68 <=> (~ (e0 = op(e4, e4)) & ~ (e0 = op(e3, e4)) & ~ (e0 = op(e2, e4)) & ~ (e0 = op(e1, e4)) & ~ (e0 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP68])])).
fof(f1426, plain, (sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1425, f843])).
fof(f843, plain, ~ sP69, inference(subsumption_resolution, [], [f230, f194])).
fof(f230, plain, (~ (e1 = op(e4, e4)) | ~ sP69), inference(cnf_transformation, [], [f90])).
fof(f90, plain, ((~ (e1 = op(e4, e4)) & ~ (e1 = op(e4, e3)) & ~ (e1 = op(e4, e2)) & ~ (e1 = op(e4, e1)) & ~ (e1 = op(e4, e0))) | ~ sP69), inference(nnf_transformation, [], [f76])).
fof(f76, plain, ((~ (e1 = op(e4, e4)) & ~ (e1 = op(e4, e3)) & ~ (e1 = op(e4, e2)) & ~ (e1 = op(e4, e1)) & ~ (e1 = op(e4, e0))) | ~ sP69), inference(usedef, [], [e76])).
fof(e76, plain, (sP69 <=> (~ (e1 = op(e4, e4)) & ~ (e1 = op(e4, e3)) & ~ (e1 = op(e4, e2)) & ~ (e1 = op(e4, e1)) & ~ (e1 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP69])])).
fof(f1425, plain, (sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1424, f842])).
fof(f842, plain, ~ sP70, inference(subsumption_resolution, [], [f225, f194])).
fof(f225, plain, (~ (e1 = op(e4, e4)) | ~ sP70), inference(cnf_transformation, [], [f89])).
fof(f89, plain, ((~ (e1 = op(e4, e4)) & ~ (e1 = op(e3, e4)) & ~ (e1 = op(e2, e4)) & ~ (e1 = op(e1, e4)) & ~ (e1 = op(e0, e4))) | ~ sP70), inference(nnf_transformation, [], [f77])).
fof(f77, plain, ((~ (e1 = op(e4, e4)) & ~ (e1 = op(e3, e4)) & ~ (e1 = op(e2, e4)) & ~ (e1 = op(e1, e4)) & ~ (e1 = op(e0, e4))) | ~ sP70), inference(usedef, [], [e77])).
fof(e77, plain, (sP70 <=> (~ (e1 = op(e4, e4)) & ~ (e1 = op(e3, e4)) & ~ (e1 = op(e2, e4)) & ~ (e1 = op(e1, e4)) & ~ (e1 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP70])])).
fof(f1424, plain, (sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1423, f841])).
fof(f841, plain, ~ sP71, inference(subsumption_resolution, [], [f219, f193])).
fof(f219, plain, (~ (e2 = op(e4, e3)) | ~ sP71), inference(cnf_transformation, [], [f88])).
fof(f88, plain, ((~ (e2 = op(e4, e4)) & ~ (e2 = op(e4, e3)) & ~ (e2 = op(e4, e2)) & ~ (e2 = op(e4, e1)) & ~ (e2 = op(e4, e0))) | ~ sP71), inference(nnf_transformation, [], [f78])).
fof(f78, plain, ((~ (e2 = op(e4, e4)) & ~ (e2 = op(e4, e3)) & ~ (e2 = op(e4, e2)) & ~ (e2 = op(e4, e1)) & ~ (e2 = op(e4, e0))) | ~ sP71), inference(usedef, [], [e78])).
fof(e78, plain, (sP71 <=> (~ (e2 = op(e4, e4)) & ~ (e2 = op(e4, e3)) & ~ (e2 = op(e4, e2)) & ~ (e2 = op(e4, e1)) & ~ (e2 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP71])])).
fof(f1423, plain, (sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1422, f839])).
fof(f839, plain, ~ sP72, inference(subsumption_resolution, [], [f214, f189])).
fof(f214, plain, (~ (e2 = op(e3, e4)) | ~ sP72), inference(cnf_transformation, [], [f87])).
fof(f87, plain, ((~ (e2 = op(e4, e4)) & ~ (e2 = op(e3, e4)) & ~ (e2 = op(e2, e4)) & ~ (e2 = op(e1, e4)) & ~ (e2 = op(e0, e4))) | ~ sP72), inference(nnf_transformation, [], [f79])).
fof(f79, plain, ((~ (e2 = op(e4, e4)) & ~ (e2 = op(e3, e4)) & ~ (e2 = op(e2, e4)) & ~ (e2 = op(e1, e4)) & ~ (e2 = op(e0, e4))) | ~ sP72), inference(usedef, [], [e79])).
fof(e79, plain, (sP72 <=> (~ (e2 = op(e4, e4)) & ~ (e2 = op(e3, e4)) & ~ (e2 = op(e2, e4)) & ~ (e2 = op(e1, e4)) & ~ (e2 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP72])])).
fof(f1422, plain, (sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1421, f837])).
fof(f837, plain, ~ sP73, inference(subsumption_resolution, [], [f207, f191])).
fof(f207, plain, (~ (e3 = op(e4, e1)) | ~ sP73), inference(cnf_transformation, [], [f86])).
fof(f86, plain, ((~ (e3 = op(e4, e4)) & ~ (e3 = op(e4, e3)) & ~ (e3 = op(e4, e2)) & ~ (e3 = op(e4, e1)) & ~ (e3 = op(e4, e0))) | ~ sP73), inference(nnf_transformation, [], [f80])).
fof(f80, plain, ((~ (e3 = op(e4, e4)) & ~ (e3 = op(e4, e3)) & ~ (e3 = op(e4, e2)) & ~ (e3 = op(e4, e1)) & ~ (e3 = op(e4, e0))) | ~ sP73), inference(usedef, [], [e80])).
fof(e80, plain, (sP73 <=> (~ (e3 = op(e4, e4)) & ~ (e3 = op(e4, e3)) & ~ (e3 = op(e4, e2)) & ~ (e3 = op(e4, e1)) & ~ (e3 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP73])])).
fof(f1421, plain, (sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1420, f833])).
fof(f833, plain, ~ sP74, inference(subsumption_resolution, [], [f203, f184])).
fof(f203, plain, (~ (e3 = op(e2, e4)) | ~ sP74), inference(cnf_transformation, [], [f85])).
fof(f85, plain, ((~ (e3 = op(e4, e4)) & ~ (e3 = op(e3, e4)) & ~ (e3 = op(e2, e4)) & ~ (e3 = op(e1, e4)) & ~ (e3 = op(e0, e4))) | ~ sP74), inference(nnf_transformation, [], [f81])).
fof(f81, plain, ((~ (e3 = op(e4, e4)) & ~ (e3 = op(e3, e4)) & ~ (e3 = op(e2, e4)) & ~ (e3 = op(e1, e4)) & ~ (e3 = op(e0, e4))) | ~ sP74), inference(usedef, [], [e81])).
fof(e81, plain, (sP74 <=> (~ (e3 = op(e4, e4)) & ~ (e3 = op(e3, e4)) & ~ (e3 = op(e2, e4)) & ~ (e3 = op(e1, e4)) & ~ (e3 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP74])])).
fof(f1420, plain, (sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1419, f830])).
fof(f830, plain, ~ sP75, inference(subsumption_resolution, [], [f618, f606])).
fof(f618, plain, (~ (e4 = op(e4, unit)) | ~ sP75), inference(definition_unfolding, [], [f196, f195])).
fof(f196, plain, (~ (e4 = op(e4, e0)) | ~ sP75), inference(cnf_transformation, [], [f84])).
fof(f84, plain, ((~ (e4 = op(e4, e4)) & ~ (e4 = op(e4, e3)) & ~ (e4 = op(e4, e2)) & ~ (e4 = op(e4, e1)) & ~ (e4 = op(e4, e0))) | ~ sP75), inference(nnf_transformation, [], [f82])).
fof(f82, plain, ((~ (e4 = op(e4, e4)) & ~ (e4 = op(e4, e3)) & ~ (e4 = op(e4, e2)) & ~ (e4 = op(e4, e1)) & ~ (e4 = op(e4, e0))) | ~ sP75), inference(usedef, [], [e82])).
fof(e82, plain, (sP75 <=> (~ (e4 = op(e4, e4)) & ~ (e4 = op(e4, e3)) & ~ (e4 = op(e4, e2)) & ~ (e4 = op(e4, e1)) & ~ (e4 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP75])])).
fof(f1419, plain, (sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f815, f613])).
fof(f815, plain, (~ (e4 = op(unit, e4)) | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(duplicate_literal_removal, [], [f814])).
fof(f814, plain, (~ (e4 = op(unit, e4)) | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | ~ (unit = op(unit, unit)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(definition_unfolding, [], [f596, f195, f195, f195, f195, f195])).
fof(f596, plain, (~ (e4 = op(e0, e4)) | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (e0 = op(e0, unit)) | ~ (e0 = op(unit, e0)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f83])).
fof(f83, plain, ((~ (e4 = op(e4, e4)) & ~ (e4 = op(e3, e4)) & ~ (e4 = op(e2, e4)) & ~ (e4 = op(e1, e4)) & ~ (e4 = op(e0, e4))) | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (e0 = op(e0, unit)) | ~ (e0 = op(unit, e0)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(definition_folding, [], [f6, e82, e81, e80, e79, e78, e77, e76, e75, e74, e73, e72, e71, e70, e69, e68, e67, e66, e65, e64, e63, e62, e61, e60, e59, e58, e57, e56, e55, e54, e53, e52, e51, e50, e49, e48, e47, e46, e45, e44, e43, e42, e41, e40, e39, e38, e37, e36, e35, e34, e33, e32, e31, e30, e29, e28, e27, e26, e25, e24, e23, e22, e21, e20, e19, e18, e17, e16, e15, e14, e13, e12, e11, e10, e9, e8, e7])).
fof(f6, plain, ((~ (e4 = op(e4, e4)) & ~ (e4 = op(e3, e4)) & ~ (e4 = op(e2, e4)) & ~ (e4 = op(e1, e4)) & ~ (e4 = op(e0, e4))) | (~ (e4 = op(e4, e4)) & ~ (e4 = op(e4, e3)) & ~ (e4 = op(e4, e2)) & ~ (e4 = op(e4, e1)) & ~ (e4 = op(e4, e0))) | (~ (e3 = op(e4, e4)) & ~ (e3 = op(e3, e4)) & ~ (e3 = op(e2, e4)) & ~ (e3 = op(e1, e4)) & ~ (e3 = op(e0, e4))) | (~ (e3 = op(e4, e4)) & ~ (e3 = op(e4, e3)) & ~ (e3 = op(e4, e2)) & ~ (e3 = op(e4, e1)) & ~ (e3 = op(e4, e0))) | (~ (e2 = op(e4, e4)) & ~ (e2 = op(e3, e4)) & ~ (e2 = op(e2, e4)) & ~ (e2 = op(e1, e4)) & ~ (e2 = op(e0, e4))) | (~ (e2 = op(e4, e4)) & ~ (e2 = op(e4, e3)) & ~ (e2 = op(e4, e2)) & ~ (e2 = op(e4, e1)) & ~ (e2 = op(e4, e0))) | (~ (e1 = op(e4, e4)) & ~ (e1 = op(e3, e4)) & ~ (e1 = op(e2, e4)) & ~ (e1 = op(e1, e4)) & ~ (e1 = op(e0, e4))) | (~ (e1 = op(e4, e4)) & ~ (e1 = op(e4, e3)) & ~ (e1 = op(e4, e2)) & ~ (e1 = op(e4, e1)) & ~ (e1 = op(e4, e0))) | (~ (e0 = op(e4, e4)) & ~ (e0 = op(e3, e4)) & ~ (e0 = op(e2, e4)) & ~ (e0 = op(e1, e4)) & ~ (e0 = op(e0, e4))) | (~ (e0 = op(e4, e4)) & ~ (e0 = op(e4, e3)) & ~ (e0 = op(e4, e2)) & ~ (e0 = op(e4, e1)) & ~ (e0 = op(e4, e0))) | (~ (e4 = op(e4, e3)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e0, e3))) | (~ (e4 = op(e3, e4)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e3, e0))) | (~ (e3 = op(e4, e3)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e0, e3))) | (~ (e3 = op(e3, e4)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e3, e0))) | (~ (e2 = op(e4, e3)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e0, e3))) | (~ (e2 = op(e3, e4)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e3, e0))) | (~ (e1 = op(e4, e3)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e0, e3))) | (~ (e1 = op(e3, e4)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e3, e0))) | (~ (e0 = op(e4, e3)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e0, e3))) | (~ (e0 = op(e3, e4)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e3, e0))) | (~ (e4 = op(e4, e2)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e0, e2))) | (~ (e4 = op(e2, e4)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e2, e0))) | (~ (e3 = op(e4, e2)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e0, e2))) | (~ (e3 = op(e2, e4)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e2, e0))) | (~ (e2 = op(e4, e2)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e0, e2))) | (~ (e2 = op(e2, e4)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e2, e0))) | (~ (e1 = op(e4, e2)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e0, e2))) | (~ (e1 = op(e2, e4)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e2, e0))) | (~ (e0 = op(e4, e2)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e0, e2))) | (~ (e0 = op(e2, e4)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e2, e0))) | (~ (e4 = op(e4, e1)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e0, e1))) | (~ (e4 = op(e1, e4)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e1, e0))) | (~ (e3 = op(e4, e1)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e0, e1))) | (~ (e3 = op(e1, e4)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e1, e0))) | (~ (e2 = op(e4, e1)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e0, e1))) | (~ (e2 = op(e1, e4)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e1, e0))) | (~ (e1 = op(e4, e1)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e0, e1))) | (~ (e1 = op(e1, e4)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e1, e0))) | (~ (e0 = op(e4, e1)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e0, e1))) | (~ (e0 = op(e1, e4)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e1, e0))) | (~ (e4 = op(e4, e0)) & ~ (e4 = op(e3, e0)) & ~ (e4 = op(e2, e0)) & ~ (e4 = op(e1, e0)) & ~ (e4 = op(e0, e0))) | (~ (e4 = op(e0, e4)) & ~ (e4 = op(e0, e3)) & ~ (e4 = op(e0, e2)) & ~ (e4 = op(e0, e1)) & ~ (e4 = op(e0, e0))) | (~ (e3 = op(e4, e0)) & ~ (e3 = op(e3, e0)) & ~ (e3 = op(e2, e0)) & ~ (e3 = op(e1, e0)) & ~ (e3 = op(e0, e0))) | (~ (e3 = op(e0, e4)) & ~ (e3 = op(e0, e3)) & ~ (e3 = op(e0, e2)) & ~ (e3 = op(e0, e1)) & ~ (e3 = op(e0, e0))) | (~ (e2 = op(e4, e0)) & ~ (e2 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e2 = op(e1, e0)) & ~ (e2 = op(e0, e0))) | (~ (e2 = op(e0, e4)) & ~ (e2 = op(e0, e3)) & ~ (e2 = op(e0, e2)) & ~ (e2 = op(e0, e1)) & ~ (e2 = op(e0, e0))) | (~ (e1 = op(e4, e0)) & ~ (e1 = op(e3, e0)) & ~ (e1 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e1 = op(e0, e0))) | (~ (e1 = op(e0, e4)) & ~ (e1 = op(e0, e3)) & ~ (e1 = op(e0, e2)) & ~ (e1 = op(e0, e1)) & ~ (e1 = op(e0, e0))) | (~ (e0 = op(e4, e0)) & ~ (e0 = op(e3, e0)) & ~ (e0 = op(e2, e0)) & ~ (e0 = op(e1, e0)) & ~ (e0 = op(e0, e0))) | (~ (e0 = op(e0, e4)) & ~ (e0 = op(e0, e3)) & ~ (e0 = op(e0, e2)) & ~ (e0 = op(e0, e1)) & ~ (e0 = op(e0, e0))) | (~ (e4 = unit) & ~ (e3 = unit) & ~ (e2 = unit) & ~ (e1 = unit) & ~ (e0 = unit)) | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (e0 = op(e0, unit)) | ~ (e0 = op(unit, e0)) | (~ (e4 = op(e4, e4)) & ~ (e3 = op(e4, e4)) & ~ (e2 = op(e4, e4)) & ~ (e1 = op(e4, e4)) & ~ (e0 = op(e4, e4))) | (~ (e4 = op(e4, e3)) & ~ (e3 = op(e4, e3)) & ~ (e2 = op(e4, e3)) & ~ (e1 = op(e4, e3)) & ~ (e0 = op(e4, e3))) | (~ (e4 = op(e4, e2)) & ~ (e3 = op(e4, e2)) & ~ (e2 = op(e4, e2)) & ~ (e1 = op(e4, e2)) & ~ (e0 = op(e4, e2))) | (~ (e4 = op(e4, e1)) & ~ (e3 = op(e4, e1)) & ~ (e2 = op(e4, e1)) & ~ (e1 = op(e4, e1)) & ~ (e0 = op(e4, e1))) | (~ (e4 = op(e4, e0)) & ~ (e3 = op(e4, e0)) & ~ (e2 = op(e4, e0)) & ~ (e1 = op(e4, e0)) & ~ (e0 = op(e4, e0))) | (~ (e4 = op(e3, e4)) & ~ (e3 = op(e3, e4)) & ~ (e2 = op(e3, e4)) & ~ (e1 = op(e3, e4)) & ~ (e0 = op(e3, e4))) | (~ (e4 = op(e3, e3)) & ~ (e3 = op(e3, e3)) & ~ (e2 = op(e3, e3)) & ~ (e1 = op(e3, e3)) & ~ (e0 = op(e3, e3))) | (~ (e4 = op(e3, e2)) & ~ (e3 = op(e3, e2)) & ~ (e2 = op(e3, e2)) & ~ (e1 = op(e3, e2)) & ~ (e0 = op(e3, e2))) | (~ (e4 = op(e3, e1)) & ~ (e3 = op(e3, e1)) & ~ (e2 = op(e3, e1)) & ~ (e1 = op(e3, e1)) & ~ (e0 = op(e3, e1))) | (~ (e4 = op(e3, e0)) & ~ (e3 = op(e3, e0)) & ~ (e2 = op(e3, e0)) & ~ (e1 = op(e3, e0)) & ~ (e0 = op(e3, e0))) | (~ (e4 = op(e2, e4)) & ~ (e3 = op(e2, e4)) & ~ (e2 = op(e2, e4)) & ~ (e1 = op(e2, e4)) & ~ (e0 = op(e2, e4))) | (~ (e4 = op(e2, e3)) & ~ (e3 = op(e2, e3)) & ~ (e2 = op(e2, e3)) & ~ (e1 = op(e2, e3)) & ~ (e0 = op(e2, e3))) | (~ (e4 = op(e2, e2)) & ~ (e3 = op(e2, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e2, e2)) & ~ (e0 = op(e2, e2))) | (~ (e4 = op(e2, e1)) & ~ (e3 = op(e2, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e2, e1)) & ~ (e0 = op(e2, e1))) | (~ (e4 = op(e2, e0)) & ~ (e3 = op(e2, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e2, e0)) & ~ (e0 = op(e2, e0))) | (~ (e4 = op(e1, e4)) & ~ (e3 = op(e1, e4)) & ~ (e2 = op(e1, e4)) & ~ (e1 = op(e1, e4)) & ~ (e0 = op(e1, e4))) | (~ (e4 = op(e1, e3)) & ~ (e3 = op(e1, e3)) & ~ (e2 = op(e1, e3)) & ~ (e1 = op(e1, e3)) & ~ (e0 = op(e1, e3))) | (~ (e4 = op(e1, e2)) & ~ (e3 = op(e1, e2)) & ~ (e2 = op(e1, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e1, e2))) | (~ (e4 = op(e1, e1)) & ~ (e3 = op(e1, e1)) & ~ (e2 = op(e1, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e1, e1))) | (~ (e4 = op(e1, e0)) & ~ (e3 = op(e1, e0)) & ~ (e2 = op(e1, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e1, e0))) | (~ (e4 = op(e0, e4)) & ~ (e3 = op(e0, e4)) & ~ (e2 = op(e0, e4)) & ~ (e1 = op(e0, e4)) & ~ (e0 = op(e0, e4))) | (~ (e4 = op(e0, e3)) & ~ (e3 = op(e0, e3)) & ~ (e2 = op(e0, e3)) & ~ (e1 = op(e0, e3)) & ~ (e0 = op(e0, e3))) | (~ (e4 = op(e0, e2)) & ~ (e3 = op(e0, e2)) & ~ (e2 = op(e0, e2)) & ~ (e1 = op(e0, e2)) & ~ (e0 = op(e0, e2))) | (~ (e4 = op(e0, e1)) & ~ (e3 = op(e0, e1)) & ~ (e2 = op(e0, e1)) & ~ (e1 = op(e0, e1)) & ~ (e0 = op(e0, e1))) | (~ (e4 = op(e0, e0)) & ~ (e3 = op(e0, e0)) & ~ (e2 = op(e0, e0)) & ~ (e1 = op(e0, e0)) & ~ (e0 = op(e0, e0))) | (((e4 = op(e4, e4)) | ~ (e4 = op(e4, e4)) | ~ (e4 = op(e4, e4))) & ((e4 = op(e4, e3)) | ~ (e4 = op(e3, e3)) | ~ (e3 = op(e4, e4))) & ((e4 = op(e4, e2)) | ~ (e4 = op(e2, e2)) | ~ (e2 = op(e4, e4))) & ((e4 = op(e4, e1)) | ~ (e4 = op(e1, e1)) | ~ (e1 = op(e4, e4))) & ((e4 = op(e4, e0)) | ~ (e4 = op(e0, e0)) | ~ (e0 = op(e4, e4))) & ((e3 = op(e3, e4)) | ~ (e3 = op(e4, e4)) | ~ (e4 = op(e3, e3))) & ((e3 = op(e3, e3)) | ~ (e3 = op(e3, e3)) | ~ (e3 = op(e3, e3))) & ((e3 = op(e3, e2)) | ~ (e3 = op(e2, e2)) | ~ (e2 = op(e3, e3))) & ((e3 = op(e3, e1)) | ~ (e3 = op(e1, e1)) | ~ (e1 = op(e3, e3))) & ((e3 = op(e3, e0)) | ~ (e3 = op(e0, e0)) | ~ (e0 = op(e3, e3))) & ((e2 = op(e2, e4)) | ~ (e2 = op(e4, e4)) | ~ (e4 = op(e2, e2))) & ((e2 = op(e2, e3)) | ~ (e2 = op(e3, e3)) | ~ (e3 = op(e2, e2))) & ((e2 = op(e2, e2)) | ~ (e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & ((e2 = op(e2, e1)) | ~ (e2 = op(e1, e1)) | ~ (e1 = op(e2, e2))) & ((e2 = op(e2, e0)) | ~ (e2 = op(e0, e0)) | ~ (e0 = op(e2, e2))) & ((e1 = op(e1, e4)) | ~ (e1 = op(e4, e4)) | ~ (e4 = op(e1, e1))) & ((e1 = op(e1, e3)) | ~ (e1 = op(e3, e3)) | ~ (e3 = op(e1, e1))) & ((e1 = op(e1, e2)) | ~ (e1 = op(e2, e2)) | ~ (e2 = op(e1, e1))) & ((e1 = op(e1, e1)) | ~ (e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & ((e1 = op(e1, e0)) | ~ (e1 = op(e0, e0)) | ~ (e0 = op(e1, e1))) & ((e0 = op(e0, e4)) | ~ (e0 = op(e4, e4)) | ~ (e4 = op(e0, e0))) & ((e0 = op(e0, e3)) | ~ (e0 = op(e3, e3)) | ~ (e3 = op(e0, e0))) & ((e0 = op(e0, e2)) | ~ (e0 = op(e2, e2)) | ~ (e2 = op(e0, e0))) & ((e0 = op(e0, e1)) | ~ (e0 = op(e1, e1)) | ~ (e1 = op(e0, e0))) & ((e0 = op(e0, e0)) | ~ (e0 = op(e0, e0)) | ~ (e0 = op(e0, e0))))), inference(ennf_transformation, [], [f5])).
fof(f5, plain, ~ (((e4 = op(e4, e4)) | (e4 = op(e3, e4)) | (e4 = op(e2, e4)) | (e4 = op(e1, e4)) | (e4 = op(e0, e4))) & ((e4 = op(e4, e4)) | (e4 = op(e4, e3)) | (e4 = op(e4, e2)) | (e4 = op(e4, e1)) | (e4 = op(e4, e0))) & ((e3 = op(e4, e4)) | (e3 = op(e3, e4)) | (e3 = op(e2, e4)) | (e3 = op(e1, e4)) | (e3 = op(e0, e4))) & ((e3 = op(e4, e4)) | (e3 = op(e4, e3)) | (e3 = op(e4, e2)) | (e3 = op(e4, e1)) | (e3 = op(e4, e0))) & ((e2 = op(e4, e4)) | (e2 = op(e3, e4)) | (e2 = op(e2, e4)) | (e2 = op(e1, e4)) | (e2 = op(e0, e4))) & ((e2 = op(e4, e4)) | (e2 = op(e4, e3)) | (e2 = op(e4, e2)) | (e2 = op(e4, e1)) | (e2 = op(e4, e0))) & ((e1 = op(e4, e4)) | (e1 = op(e3, e4)) | (e1 = op(e2, e4)) | (e1 = op(e1, e4)) | (e1 = op(e0, e4))) & ((e1 = op(e4, e4)) | (e1 = op(e4, e3)) | (e1 = op(e4, e2)) | (e1 = op(e4, e1)) | (e1 = op(e4, e0))) & ((e0 = op(e4, e4)) | (e0 = op(e3, e4)) | (e0 = op(e2, e4)) | (e0 = op(e1, e4)) | (e0 = op(e0, e4))) & ((e0 = op(e4, e4)) | (e0 = op(e4, e3)) | (e0 = op(e4, e2)) | (e0 = op(e4, e1)) | (e0 = op(e4, e0))) & ((e4 = op(e4, e3)) | (e4 = op(e3, e3)) | (e4 = op(e2, e3)) | (e4 = op(e1, e3)) | (e4 = op(e0, e3))) & ((e4 = op(e3, e4)) | (e4 = op(e3, e3)) | (e4 = op(e3, e2)) | (e4 = op(e3, e1)) | (e4 = op(e3, e0))) & ((e3 = op(e4, e3)) | (e3 = op(e3, e3)) | (e3 = op(e2, e3)) | (e3 = op(e1, e3)) | (e3 = op(e0, e3))) & ((e3 = op(e3, e4)) | (e3 = op(e3, e3)) | (e3 = op(e3, e2)) | (e3 = op(e3, e1)) | (e3 = op(e3, e0))) & ((e2 = op(e4, e3)) | (e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(e0, e3))) & ((e2 = op(e3, e4)) | (e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, e0))) & ((e1 = op(e4, e3)) | (e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(e0, e3))) & ((e1 = op(e3, e4)) | (e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, e0))) & ((e0 = op(e4, e3)) | (e0 = op(e3, e3)) | (e0 = op(e2, e3)) | (e0 = op(e1, e3)) | (e0 = op(e0, e3))) & ((e0 = op(e3, e4)) | (e0 = op(e3, e3)) | (e0 = op(e3, e2)) | (e0 = op(e3, e1)) | (e0 = op(e3, e0))) & ((e4 = op(e4, e2)) | (e4 = op(e3, e2)) | (e4 = op(e2, e2)) | (e4 = op(e1, e2)) | (e4 = op(e0, e2))) & ((e4 = op(e2, e4)) | (e4 = op(e2, e3)) | (e4 = op(e2, e2)) | (e4 = op(e2, e1)) | (e4 = op(e2, e0))) & ((e3 = op(e4, e2)) | (e3 = op(e3, e2)) | (e3 = op(e2, e2)) | (e3 = op(e1, e2)) | (e3 = op(e0, e2))) & ((e3 = op(e2, e4)) | (e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, e0))) & ((e2 = op(e4, e2)) | (e2 = op(e3, e2)) | (e2 = op(e2, e2)) | (e2 = op(e1, e2)) | (e2 = op(e0, e2))) & ((e2 = op(e2, e4)) | (e2 = op(e2, e3)) | (e2 = op(e2, e2)) | (e2 = op(e2, e1)) | (e2 = op(e2, e0))) & ((e1 = op(e4, e2)) | (e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(e0, e2))) & ((e1 = op(e2, e4)) | (e1 = op(e2, e3)) | (e1 = op(e2, e2)) | (e1 = op(e2, e1)) | (e1 = op(e2, e0))) & ((e0 = op(e4, e2)) | (e0 = op(e3, e2)) | (e0 = op(e2, e2)) | (e0 = op(e1, e2)) | (e0 = op(e0, e2))) & ((e0 = op(e2, e4)) | (e0 = op(e2, e3)) | (e0 = op(e2, e2)) | (e0 = op(e2, e1)) | (e0 = op(e2, e0))) & ((e4 = op(e4, e1)) | (e4 = op(e3, e1)) | (e4 = op(e2, e1)) | (e4 = op(e1, e1)) | (e4 = op(e0, e1))) & ((e4 = op(e1, e4)) | (e4 = op(e1, e3)) | (e4 = op(e1, e2)) | (e4 = op(e1, e1)) | (e4 = op(e1, e0))) & ((e3 = op(e4, e1)) | (e3 = op(e3, e1)) | (e3 = op(e2, e1)) | (e3 = op(e1, e1)) | (e3 = op(e0, e1))) & ((e3 = op(e1, e4)) | (e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, e0))) & ((e2 = op(e4, e1)) | (e2 = op(e3, e1)) | (e2 = op(e2, e1)) | (e2 = op(e1, e1)) | (e2 = op(e0, e1))) & ((e2 = op(e1, e4)) | (e2 = op(e1, e3)) | (e2 = op(e1, e2)) | (e2 = op(e1, e1)) | (e2 = op(e1, e0))) & ((e1 = op(e4, e1)) | (e1 = op(e3, e1)) | (e1 = op(e2, e1)) | (e1 = op(e1, e1)) | (e1 = op(e0, e1))) & ((e1 = op(e1, e4)) | (e1 = op(e1, e3)) | (e1 = op(e1, e2)) | (e1 = op(e1, e1)) | (e1 = op(e1, e0))) & ((e0 = op(e4, e1)) | (e0 = op(e3, e1)) | (e0 = op(e2, e1)) | (e0 = op(e1, e1)) | (e0 = op(e0, e1))) & ((e0 = op(e1, e4)) | (e0 = op(e1, e3)) | (e0 = op(e1, e2)) | (e0 = op(e1, e1)) | (e0 = op(e1, e0))) & ((e4 = op(e4, e0)) | (e4 = op(e3, e0)) | (e4 = op(e2, e0)) | (e4 = op(e1, e0)) | (e4 = op(e0, e0))) & ((e4 = op(e0, e4)) | (e4 = op(e0, e3)) | (e4 = op(e0, e2)) | (e4 = op(e0, e1)) | (e4 = op(e0, e0))) & ((e3 = op(e4, e0)) | (e3 = op(e3, e0)) | (e3 = op(e2, e0)) | (e3 = op(e1, e0)) | (e3 = op(e0, e0))) & ((e3 = op(e0, e4)) | (e3 = op(e0, e3)) | (e3 = op(e0, e2)) | (e3 = op(e0, e1)) | (e3 = op(e0, e0))) & ((e2 = op(e4, e0)) | (e2 = op(e3, e0)) | (e2 = op(e2, e0)) | (e2 = op(e1, e0)) | (e2 = op(e0, e0))) & ((e2 = op(e0, e4)) | (e2 = op(e0, e3)) | (e2 = op(e0, e2)) | (e2 = op(e0, e1)) | (e2 = op(e0, e0))) & ((e1 = op(e4, e0)) | (e1 = op(e3, e0)) | (e1 = op(e2, e0)) | (e1 = op(e1, e0)) | (e1 = op(e0, e0))) & ((e1 = op(e0, e4)) | (e1 = op(e0, e3)) | (e1 = op(e0, e2)) | (e1 = op(e0, e1)) | (e1 = op(e0, e0))) & ((e0 = op(e4, e0)) | (e0 = op(e3, e0)) | (e0 = op(e2, e0)) | (e0 = op(e1, e0)) | (e0 = op(e0, e0))) & ((e0 = op(e0, e4)) | (e0 = op(e0, e3)) | (e0 = op(e0, e2)) | (e0 = op(e0, e1)) | (e0 = op(e0, e0))) & ((e4 = unit) | (e3 = unit) | (e2 = unit) | (e1 = unit) | (e0 = unit)) & (e4 = op(e4, unit)) & (e4 = op(unit, e4)) & (e3 = op(e3, unit)) & (e3 = op(unit, e3)) & (e2 = op(e2, unit)) & (e2 = op(unit, e2)) & (e1 = op(e1, unit)) & (e1 = op(unit, e1)) & (e0 = op(e0, unit)) & (e0 = op(unit, e0)) & ((e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | (e1 = op(e4, e4)) | (e0 = op(e4, e4))) & ((e4 = op(e4, e3)) | (e3 = op(e4, e3)) | (e2 = op(e4, e3)) | (e1 = op(e4, e3)) | (e0 = op(e4, e3))) & ((e4 = op(e4, e2)) | (e3 = op(e4, e2)) | (e2 = op(e4, e2)) | (e1 = op(e4, e2)) | (e0 = op(e4, e2))) & ((e4 = op(e4, e1)) | (e3 = op(e4, e1)) | (e2 = op(e4, e1)) | (e1 = op(e4, e1)) | (e0 = op(e4, e1))) & ((e4 = op(e4, e0)) | (e3 = op(e4, e0)) | (e2 = op(e4, e0)) | (e1 = op(e4, e0)) | (e0 = op(e4, e0))) & ((e4 = op(e3, e4)) | (e3 = op(e3, e4)) | (e2 = op(e3, e4)) | (e1 = op(e3, e4)) | (e0 = op(e3, e4))) & ((e4 = op(e3, e3)) | (e3 = op(e3, e3)) | (e2 = op(e3, e3)) | (e1 = op(e3, e3)) | (e0 = op(e3, e3))) & ((e4 = op(e3, e2)) | (e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (e0 = op(e3, e2))) & ((e4 = op(e3, e1)) | (e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (e0 = op(e3, e1))) & ((e4 = op(e3, e0)) | (e3 = op(e3, e0)) | (e2 = op(e3, e0)) | (e1 = op(e3, e0)) | (e0 = op(e3, e0))) & ((e4 = op(e2, e4)) | (e3 = op(e2, e4)) | (e2 = op(e2, e4)) | (e1 = op(e2, e4)) | (e0 = op(e2, e4))) & ((e4 = op(e2, e3)) | (e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))) & ((e4 = op(e2, e2)) | (e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (e0 = op(e2, e2))) & ((e4 = op(e2, e1)) | (e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))) & ((e4 = op(e2, e0)) | (e3 = op(e2, e0)) | (e2 = op(e2, e0)) | (e1 = op(e2, e0)) | (e0 = op(e2, e0))) & ((e4 = op(e1, e4)) | (e3 = op(e1, e4)) | (e2 = op(e1, e4)) | (e1 = op(e1, e4)) | (e0 = op(e1, e4))) & ((e4 = op(e1, e3)) | (e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (e0 = op(e1, e3))) & ((e4 = op(e1, e2)) | (e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (e0 = op(e1, e2))) & ((e4 = op(e1, e1)) | (e3 = op(e1, e1)) | (e2 = op(e1, e1)) | (e1 = op(e1, e1)) | (e0 = op(e1, e1))) & ((e4 = op(e1, e0)) | (e3 = op(e1, e0)) | (e2 = op(e1, e0)) | (e1 = op(e1, e0)) | (e0 = op(e1, e0))) & ((e4 = op(e0, e4)) | (e3 = op(e0, e4)) | (e2 = op(e0, e4)) | (e1 = op(e0, e4)) | (e0 = op(e0, e4))) & ((e4 = op(e0, e3)) | (e3 = op(e0, e3)) | (e2 = op(e0, e3)) | (e1 = op(e0, e3)) | (e0 = op(e0, e3))) & ((e4 = op(e0, e2)) | (e3 = op(e0, e2)) | (e2 = op(e0, e2)) | (e1 = op(e0, e2)) | (e0 = op(e0, e2))) & ((e4 = op(e0, e1)) | (e3 = op(e0, e1)) | (e2 = op(e0, e1)) | (e1 = op(e0, e1)) | (e0 = op(e0, e1))) & ((e4 = op(e0, e0)) | (e3 = op(e0, e0)) | (e2 = op(e0, e0)) | (e1 = op(e0, e0)) | (e0 = op(e0, e0))) & ((~ (e4 = op(e4, e4)) & (e4 = op(e4, e4)) & (e4 = op(e4, e4))) | (~ (e4 = op(e4, e3)) & (e4 = op(e3, e3)) & (e3 = op(e4, e4))) | (~ (e4 = op(e4, e2)) & (e4 = op(e2, e2)) & (e2 = op(e4, e4))) | (~ (e4 = op(e4, e1)) & (e4 = op(e1, e1)) & (e1 = op(e4, e4))) | (~ (e4 = op(e4, e0)) & (e4 = op(e0, e0)) & (e0 = op(e4, e4))) | (~ (e3 = op(e3, e4)) & (e3 = op(e4, e4)) & (e4 = op(e3, e3))) | (~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | (~ (e3 = op(e3, e2)) & (e3 = op(e2, e2)) & (e2 = op(e3, e3))) | (~ (e3 = op(e3, e1)) & (e3 = op(e1, e1)) & (e1 = op(e3, e3))) | (~ (e3 = op(e3, e0)) & (e3 = op(e0, e0)) & (e0 = op(e3, e3))) | (~ (e2 = op(e2, e4)) & (e2 = op(e4, e4)) & (e4 = op(e2, e2))) | (~ (e2 = op(e2, e3)) & (e2 = op(e3, e3)) & (e3 = op(e2, e2))) | (~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | (~ (e2 = op(e2, e1)) & (e2 = op(e1, e1)) & (e1 = op(e2, e2))) | (~ (e2 = op(e2, e0)) & (e2 = op(e0, e0)) & (e0 = op(e2, e2))) | (~ (e1 = op(e1, e4)) & (e1 = op(e4, e4)) & (e4 = op(e1, e1))) | (~ (e1 = op(e1, e3)) & (e1 = op(e3, e3)) & (e3 = op(e1, e1))) | (~ (e1 = op(e1, e2)) & (e1 = op(e2, e2)) & (e2 = op(e1, e1))) | (~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | (~ (e1 = op(e1, e0)) & (e1 = op(e0, e0)) & (e0 = op(e1, e1))) | (~ (e0 = op(e0, e4)) & (e0 = op(e4, e4)) & (e4 = op(e0, e0))) | (~ (e0 = op(e0, e3)) & (e0 = op(e3, e3)) & (e3 = op(e0, e0))) | (~ (e0 = op(e0, e2)) & (e0 = op(e2, e2)) & (e2 = op(e0, e0))) | (~ (e0 = op(e0, e1)) & (e0 = op(e1, e1)) & (e1 = op(e0, e0))) | (~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0))))), inference(negated_conjecture, [], [f4])).
fof(f4, plain, ~ (((e4 = op(e4, e4)) | (e4 = op(e3, e4)) | (e4 = op(e2, e4)) | (e4 = op(e1, e4)) | (e4 = op(e0, e4))) & ((e4 = op(e4, e4)) | (e4 = op(e4, e3)) | (e4 = op(e4, e2)) | (e4 = op(e4, e1)) | (e4 = op(e4, e0))) & ((e3 = op(e4, e4)) | (e3 = op(e3, e4)) | (e3 = op(e2, e4)) | (e3 = op(e1, e4)) | (e3 = op(e0, e4))) & ((e3 = op(e4, e4)) | (e3 = op(e4, e3)) | (e3 = op(e4, e2)) | (e3 = op(e4, e1)) | (e3 = op(e4, e0))) & ((e2 = op(e4, e4)) | (e2 = op(e3, e4)) | (e2 = op(e2, e4)) | (e2 = op(e1, e4)) | (e2 = op(e0, e4))) & ((e2 = op(e4, e4)) | (e2 = op(e4, e3)) | (e2 = op(e4, e2)) | (e2 = op(e4, e1)) | (e2 = op(e4, e0))) & ((e1 = op(e4, e4)) | (e1 = op(e3, e4)) | (e1 = op(e2, e4)) | (e1 = op(e1, e4)) | (e1 = op(e0, e4))) & ((e1 = op(e4, e4)) | (e1 = op(e4, e3)) | (e1 = op(e4, e2)) | (e1 = op(e4, e1)) | (e1 = op(e4, e0))) & ((e0 = op(e4, e4)) | (e0 = op(e3, e4)) | (e0 = op(e2, e4)) | (e0 = op(e1, e4)) | (e0 = op(e0, e4))) & ((e0 = op(e4, e4)) | (e0 = op(e4, e3)) | (e0 = op(e4, e2)) | (e0 = op(e4, e1)) | (e0 = op(e4, e0))) & ((e4 = op(e4, e3)) | (e4 = op(e3, e3)) | (e4 = op(e2, e3)) | (e4 = op(e1, e3)) | (e4 = op(e0, e3))) & ((e4 = op(e3, e4)) | (e4 = op(e3, e3)) | (e4 = op(e3, e2)) | (e4 = op(e3, e1)) | (e4 = op(e3, e0))) & ((e3 = op(e4, e3)) | (e3 = op(e3, e3)) | (e3 = op(e2, e3)) | (e3 = op(e1, e3)) | (e3 = op(e0, e3))) & ((e3 = op(e3, e4)) | (e3 = op(e3, e3)) | (e3 = op(e3, e2)) | (e3 = op(e3, e1)) | (e3 = op(e3, e0))) & ((e2 = op(e4, e3)) | (e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(e0, e3))) & ((e2 = op(e3, e4)) | (e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, e0))) & ((e1 = op(e4, e3)) | (e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(e0, e3))) & ((e1 = op(e3, e4)) | (e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, e0))) & ((e0 = op(e4, e3)) | (e0 = op(e3, e3)) | (e0 = op(e2, e3)) | (e0 = op(e1, e3)) | (e0 = op(e0, e3))) & ((e0 = op(e3, e4)) | (e0 = op(e3, e3)) | (e0 = op(e3, e2)) | (e0 = op(e3, e1)) | (e0 = op(e3, e0))) & ((e4 = op(e4, e2)) | (e4 = op(e3, e2)) | (e4 = op(e2, e2)) | (e4 = op(e1, e2)) | (e4 = op(e0, e2))) & ((e4 = op(e2, e4)) | (e4 = op(e2, e3)) | (e4 = op(e2, e2)) | (e4 = op(e2, e1)) | (e4 = op(e2, e0))) & ((e3 = op(e4, e2)) | (e3 = op(e3, e2)) | (e3 = op(e2, e2)) | (e3 = op(e1, e2)) | (e3 = op(e0, e2))) & ((e3 = op(e2, e4)) | (e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, e0))) & ((e2 = op(e4, e2)) | (e2 = op(e3, e2)) | (e2 = op(e2, e2)) | (e2 = op(e1, e2)) | (e2 = op(e0, e2))) & ((e2 = op(e2, e4)) | (e2 = op(e2, e3)) | (e2 = op(e2, e2)) | (e2 = op(e2, e1)) | (e2 = op(e2, e0))) & ((e1 = op(e4, e2)) | (e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(e0, e2))) & ((e1 = op(e2, e4)) | (e1 = op(e2, e3)) | (e1 = op(e2, e2)) | (e1 = op(e2, e1)) | (e1 = op(e2, e0))) & ((e0 = op(e4, e2)) | (e0 = op(e3, e2)) | (e0 = op(e2, e2)) | (e0 = op(e1, e2)) | (e0 = op(e0, e2))) & ((e0 = op(e2, e4)) | (e0 = op(e2, e3)) | (e0 = op(e2, e2)) | (e0 = op(e2, e1)) | (e0 = op(e2, e0))) & ((e4 = op(e4, e1)) | (e4 = op(e3, e1)) | (e4 = op(e2, e1)) | (e4 = op(e1, e1)) | (e4 = op(e0, e1))) & ((e4 = op(e1, e4)) | (e4 = op(e1, e3)) | (e4 = op(e1, e2)) | (e4 = op(e1, e1)) | (e4 = op(e1, e0))) & ((e3 = op(e4, e1)) | (e3 = op(e3, e1)) | (e3 = op(e2, e1)) | (e3 = op(e1, e1)) | (e3 = op(e0, e1))) & ((e3 = op(e1, e4)) | (e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, e0))) & ((e2 = op(e4, e1)) | (e2 = op(e3, e1)) | (e2 = op(e2, e1)) | (e2 = op(e1, e1)) | (e2 = op(e0, e1))) & ((e2 = op(e1, e4)) | (e2 = op(e1, e3)) | (e2 = op(e1, e2)) | (e2 = op(e1, e1)) | (e2 = op(e1, e0))) & ((e1 = op(e4, e1)) | (e1 = op(e3, e1)) | (e1 = op(e2, e1)) | (e1 = op(e1, e1)) | (e1 = op(e0, e1))) & ((e1 = op(e1, e4)) | (e1 = op(e1, e3)) | (e1 = op(e1, e2)) | (e1 = op(e1, e1)) | (e1 = op(e1, e0))) & ((e0 = op(e4, e1)) | (e0 = op(e3, e1)) | (e0 = op(e2, e1)) | (e0 = op(e1, e1)) | (e0 = op(e0, e1))) & ((e0 = op(e1, e4)) | (e0 = op(e1, e3)) | (e0 = op(e1, e2)) | (e0 = op(e1, e1)) | (e0 = op(e1, e0))) & ((e4 = op(e4, e0)) | (e4 = op(e3, e0)) | (e4 = op(e2, e0)) | (e4 = op(e1, e0)) | (e4 = op(e0, e0))) & ((e4 = op(e0, e4)) | (e4 = op(e0, e3)) | (e4 = op(e0, e2)) | (e4 = op(e0, e1)) | (e4 = op(e0, e0))) & ((e3 = op(e4, e0)) | (e3 = op(e3, e0)) | (e3 = op(e2, e0)) | (e3 = op(e1, e0)) | (e3 = op(e0, e0))) & ((e3 = op(e0, e4)) | (e3 = op(e0, e3)) | (e3 = op(e0, e2)) | (e3 = op(e0, e1)) | (e3 = op(e0, e0))) & ((e2 = op(e4, e0)) | (e2 = op(e3, e0)) | (e2 = op(e2, e0)) | (e2 = op(e1, e0)) | (e2 = op(e0, e0))) & ((e2 = op(e0, e4)) | (e2 = op(e0, e3)) | (e2 = op(e0, e2)) | (e2 = op(e0, e1)) | (e2 = op(e0, e0))) & ((e1 = op(e4, e0)) | (e1 = op(e3, e0)) | (e1 = op(e2, e0)) | (e1 = op(e1, e0)) | (e1 = op(e0, e0))) & ((e1 = op(e0, e4)) | (e1 = op(e0, e3)) | (e1 = op(e0, e2)) | (e1 = op(e0, e1)) | (e1 = op(e0, e0))) & ((e0 = op(e4, e0)) | (e0 = op(e3, e0)) | (e0 = op(e2, e0)) | (e0 = op(e1, e0)) | (e0 = op(e0, e0))) & ((e0 = op(e0, e4)) | (e0 = op(e0, e3)) | (e0 = op(e0, e2)) | (e0 = op(e0, e1)) | (e0 = op(e0, e0))) & ((e4 = unit) | (e3 = unit) | (e2 = unit) | (e1 = unit) | (e0 = unit)) & (e4 = op(e4, unit)) & (e4 = op(unit, e4)) & (e3 = op(e3, unit)) & (e3 = op(unit, e3)) & (e2 = op(e2, unit)) & (e2 = op(unit, e2)) & (e1 = op(e1, unit)) & (e1 = op(unit, e1)) & (e0 = op(e0, unit)) & (e0 = op(unit, e0)) & ((e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | (e1 = op(e4, e4)) | (e0 = op(e4, e4))) & ((e4 = op(e4, e3)) | (e3 = op(e4, e3)) | (e2 = op(e4, e3)) | (e1 = op(e4, e3)) | (e0 = op(e4, e3))) & ((e4 = op(e4, e2)) | (e3 = op(e4, e2)) | (e2 = op(e4, e2)) | (e1 = op(e4, e2)) | (e0 = op(e4, e2))) & ((e4 = op(e4, e1)) | (e3 = op(e4, e1)) | (e2 = op(e4, e1)) | (e1 = op(e4, e1)) | (e0 = op(e4, e1))) & ((e4 = op(e4, e0)) | (e3 = op(e4, e0)) | (e2 = op(e4, e0)) | (e1 = op(e4, e0)) | (e0 = op(e4, e0))) & ((e4 = op(e3, e4)) | (e3 = op(e3, e4)) | (e2 = op(e3, e4)) | (e1 = op(e3, e4)) | (e0 = op(e3, e4))) & ((e4 = op(e3, e3)) | (e3 = op(e3, e3)) | (e2 = op(e3, e3)) | (e1 = op(e3, e3)) | (e0 = op(e3, e3))) & ((e4 = op(e3, e2)) | (e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (e0 = op(e3, e2))) & ((e4 = op(e3, e1)) | (e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (e0 = op(e3, e1))) & ((e4 = op(e3, e0)) | (e3 = op(e3, e0)) | (e2 = op(e3, e0)) | (e1 = op(e3, e0)) | (e0 = op(e3, e0))) & ((e4 = op(e2, e4)) | (e3 = op(e2, e4)) | (e2 = op(e2, e4)) | (e1 = op(e2, e4)) | (e0 = op(e2, e4))) & ((e4 = op(e2, e3)) | (e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))) & ((e4 = op(e2, e2)) | (e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (e0 = op(e2, e2))) & ((e4 = op(e2, e1)) | (e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))) & ((e4 = op(e2, e0)) | (e3 = op(e2, e0)) | (e2 = op(e2, e0)) | (e1 = op(e2, e0)) | (e0 = op(e2, e0))) & ((e4 = op(e1, e4)) | (e3 = op(e1, e4)) | (e2 = op(e1, e4)) | (e1 = op(e1, e4)) | (e0 = op(e1, e4))) & ((e4 = op(e1, e3)) | (e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (e0 = op(e1, e3))) & ((e4 = op(e1, e2)) | (e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (e0 = op(e1, e2))) & ((e4 = op(e1, e1)) | (e3 = op(e1, e1)) | (e2 = op(e1, e1)) | (e1 = op(e1, e1)) | (e0 = op(e1, e1))) & ((e4 = op(e1, e0)) | (e3 = op(e1, e0)) | (e2 = op(e1, e0)) | (e1 = op(e1, e0)) | (e0 = op(e1, e0))) & ((e4 = op(e0, e4)) | (e3 = op(e0, e4)) | (e2 = op(e0, e4)) | (e1 = op(e0, e4)) | (e0 = op(e0, e4))) & ((e4 = op(e0, e3)) | (e3 = op(e0, e3)) | (e2 = op(e0, e3)) | (e1 = op(e0, e3)) | (e0 = op(e0, e3))) & ((e4 = op(e0, e2)) | (e3 = op(e0, e2)) | (e2 = op(e0, e2)) | (e1 = op(e0, e2)) | (e0 = op(e0, e2))) & ((e4 = op(e0, e1)) | (e3 = op(e0, e1)) | (e2 = op(e0, e1)) | (e1 = op(e0, e1)) | (e0 = op(e0, e1))) & ((e4 = op(e0, e0)) | (e3 = op(e0, e0)) | (e2 = op(e0, e0)) | (e1 = op(e0, e0)) | (e0 = op(e0, e0))) & ((~ (e4 = op(e4, e4)) & (e4 = op(e4, e4)) & (e4 = op(e4, e4))) | (~ (e4 = op(e4, e3)) & (e4 = op(e3, e3)) & (e3 = op(e4, e4))) | (~ (e4 = op(e4, e2)) & (e4 = op(e2, e2)) & (e2 = op(e4, e4))) | (~ (e4 = op(e4, e1)) & (e4 = op(e1, e1)) & (e1 = op(e4, e4))) | (~ (e4 = op(e4, e0)) & (e4 = op(e0, e0)) & (e0 = op(e4, e4))) | (~ (e3 = op(e3, e4)) & (e3 = op(e4, e4)) & (e4 = op(e3, e3))) | (~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | (~ (e3 = op(e3, e2)) & (e3 = op(e2, e2)) & (e2 = op(e3, e3))) | (~ (e3 = op(e3, e1)) & (e3 = op(e1, e1)) & (e1 = op(e3, e3))) | (~ (e3 = op(e3, e0)) & (e3 = op(e0, e0)) & (e0 = op(e3, e3))) | (~ (e2 = op(e2, e4)) & (e2 = op(e4, e4)) & (e4 = op(e2, e2))) | (~ (e2 = op(e2, e3)) & (e2 = op(e3, e3)) & (e3 = op(e2, e2))) | (~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | (~ (e2 = op(e2, e1)) & (e2 = op(e1, e1)) & (e1 = op(e2, e2))) | (~ (e2 = op(e2, e0)) & (e2 = op(e0, e0)) & (e0 = op(e2, e2))) | (~ (e1 = op(e1, e4)) & (e1 = op(e4, e4)) & (e4 = op(e1, e1))) | (~ (e1 = op(e1, e3)) & (e1 = op(e3, e3)) & (e3 = op(e1, e1))) | (~ (e1 = op(e1, e2)) & (e1 = op(e2, e2)) & (e2 = op(e1, e1))) | (~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | (~ (e1 = op(e1, e0)) & (e1 = op(e0, e0)) & (e0 = op(e1, e1))) | (~ (e0 = op(e0, e4)) & (e0 = op(e4, e4)) & (e4 = op(e0, e0))) | (~ (e0 = op(e0, e3)) & (e0 = op(e3, e3)) & (e3 = op(e0, e0))) | (~ (e0 = op(e0, e2)) & (e0 = op(e2, e2)) & (e2 = op(e0, e0))) | (~ (e0 = op(e0, e1)) & (e0 = op(e1, e1)) & (e1 = op(e0, e0))) | (~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG094+1.p', co1)).