fof(f12841, plain, $false, inference(avatar_sat_refutation, [], [f950, f1034, f1055, f1097, f1118, f1181, f1244, f1391, f1412, f1433, f1436, f1439, f1445, f1450, f1451, f1452, f1467, f1468, f1472, f1475, f1480, f1546, f1588, f1630, f1651, f1672, f1693, f1714, f1840, f2008, f2029, f2032, f2035, f2039, f2040, f2041, f2046, f2047, f2056, f2059, f2060, f2063, f2064, f2065, f2068, f2071, f2072, f2076, f2086, f2093, f2099, f2107, f2113, f2121, f2128, f2134, f2140, f2141, f2142, f2148, f2155, f2161, f2163, f2170, f2175, f2177, f2184, f2191, f2196, f2197, f2203, f2205, f2212, f2218, f2225, f2226, f2231, f2233, f2239, f2245, f2247, f2248, f2250, f2257, f2264, f2270, f2278, f2284, f2292, f2299, f2305, f2311, f2312, f2319, f2326, f2332, f2334, f2341, f2347, f2348, f2355, f2362, f2367, f2368, f2376, f2383, f2389, f2397, f2402, f2404, f2410, f2416, f2418, f2419, f2421, f2423, f2424, f2426, f2427, f2505, f2527, f2546, f2572, f2592, f2607, f2647, f2668, f2688, f2708, f2723, f2738, f2743, f2759, f2764, f2784, f2824, f2839, f2986, f3208, f3367, f3583, f3602, f3761, f3835, f3893, f3913, f3939, f3940, f3941, f3951, f3958, f4006, f4011, f4017, f4020, f4048, f4051, f4058, f4084, f4087, f4148, f4151, f4240, f4252, f4313, f4321, f4352, f4355, f4357, f4360, f4369, f4371, f4378, f4380, f4430, f4432, f4438, f4446, f4466, f4468, f4470, f4481, f4575, f4580, f4599, f4601, f4613, f4619, f4635, f4656, f4672, f4697, f4708, f4732, f4741, f4784, f4786, f4799, f4801, f4824, f4831, f4865, f4874, f4903, f4928, f4945, f4964, f4965, f4973, f4974, f4975, f5008, f5014, f5023, f5028, f5058, f5063, f5071, f5112, f5169, f5180, f5187, f5197, f5203, f5238, f5265, f5281, f5326, f5329, f5356, f5393, f5397, f5401, f5416, f5433, f5437, f5460, f5472, f5497, f5523, f5528, f5537, f5613, f5640, f5646, f5657, f5659, f5671, f5680, f5711, f5741, f5764, f5794, f5799, f5826, f5827, f5897, f5951, f5973, f5980, f5999, f6033, f6064, f6066, f6089, f6113, f6134, f6172, f6194, f6217, f6320, f6340, f6360, f6364, f6377, f6397, f6414, f6432, f6452, f6459, f6472, f6481, f6483, f6492, f6503, f6507, f6510, f6521, f6527, f6537, f6547, f6556, f6566, f6575, f6585, f6595, f6604, f6614, f6623, f6633, f6643, f6652, f6661, f6670, f6675, f6680, f6689, f6692, f6697, f6700, f6708, f6730, f6733, f6744, f6771, f6797, f6803, f6813, f6822, f6832, f6841, f6851, f6861, f6870, f6880, f6889, f6899, f6909, f6925, f6933, f6938, f6944, f6952, f6955, f6963, f6966, f6974, f6980, f7014, f7046, f7053, f7057, f7085, f7097, f7107, f7217, f7322, f7324, f7326, f7334, f7353, f7360, f7425, f7430, f7446, f7475, f7497, f7555, f7570, f7606, f7746, f7758, f7765, f7831, f7850, f7901, f7941, f7964, f7998, f8028, f8053, f8065, f8080, f8085, f8094, f8098, f8102, f8111, f8116, f8126, f8131, f8141, f8144, f8167, f8177, f8182, f8185, f8235, f8281, f8308, f8319, f8324, f8337, f8342, f8354, f8359, f8372, f8378, f8392, f8404, f8414, f8429, f8442, f8452, f8458, f8475, f8486, f8503, f8531, f8533, f8536, f8573, f8624, f8886, f8887, f8888, f8889, f8890, f8891, f8892, f8893, f8894, f8915, f8916, f8950, f8959, f8993, f9039, f9059, f9095, f9108, f9115, f9175, f9251, f9257, f9263, f9277, f9280, f9341, f9475, f9538, f9543, f9568, f9651, f9695, f9718, f9775, f9791, f9792, f9793, f9794, f9795, f9796, f9799, f9800, f9801, f9817, f9830, f9863, f9874, f9885, f9894, f10007, f10018, f10036, f10041, f10293, f10296, f10317, f10378, f10389, f10407, f10412, f10425, f10429, f10440, f10445, f10455, f10460, f10470, f10474, f10484, f10489, f10500, f10631, f10633, f10635, f10699, f10755, f10758, f10771, f10772, f10773, f10774, f10776, f10777, f10778, f10779, f10780, f10795, f11001, f11136, f11142, f11152, f11161, f11171, f11180, f11190, f11200, f11209, f11219, f11228, f11238, f11247, f11256, f11265, f11270, f11275, f11284, f11287, f11292, f11297, f11306, f11467, f11473, f11610, f11620, f11696, f11699, f11708, f11713, f11726, f11731, f11744, f11749, f11762, f11765, f11774, f11779, f11791, f11797, f11811, f11817, f11831, f11837, f11851, f11857, f11870, f12014, f12017, f12075, f12077, f12127, f12128, f12129, f12130, f12132, f12133, f12134, f12135, f12136, f12166, f12245, f12257, f12274, f12288, f12304, f12320, f12369, f12373, f12513, f12535, f12602, f12605, f12607, f12612, f12627, f12817, f12831, f12840])).
fof(f12840, plain, (~ spl68_35 | spl68_388 | ~ spl68_432), inference(avatar_contradiction_clause, [], [f12839])).
fof(f12839, plain, ($false | (~ spl68_35 | spl68_388 | ~ spl68_432)), inference(subsumption_resolution, [], [f12838, f3225])).
fof(f3225, plain, ((e24 = h4(e11)) | ~ spl68_432), inference(avatar_component_clause, [], [f3224])).
fof(f3224, plain, (spl68_432 <=> (e24 = h4(e11))), introduced(avatar_definition, [new_symbols(naming, [spl68_432])])).
fof(f12838, plain, (~ (e24 = h4(e11)) | (~ spl68_35 | spl68_388)), inference(forward_demodulation, [], [f12837, f758])).
fof(f758, plain, (e24 = h5(e14)), inference(cnf_transformation, [], [f20])).
fof(f20, plain, ((op2(e24, op2(e24, e24)) = h5(e13)) & (op2(op2(e24, op2(e24, e24)), e24) = h5(e12)) & (op2(e24, e24) = h5(e11)) & (op2(op2(e24, e24), e24) = h5(e10)) & (e24 = h5(e14))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG052+1.p', ax20)).
fof(f12837, plain, (~ (h4(e11) = h5(e14)) | (~ spl68_35 | spl68_388)), inference(forward_demodulation, [], [f2929, f1033])).
fof(f1033, plain, ((e14 = op1(e13, e13)) | ~ spl68_35), inference(avatar_component_clause, [], [f1031])).
fof(f1031, plain, (spl68_35 <=> (e14 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl68_35])])).
fof(f2929, plain, (~ (h4(e11) = h5(op1(e13, e13))) | spl68_388), inference(avatar_component_clause, [], [f2927])).
fof(f2927, plain, (spl68_388 <=> (h4(e11) = h5(op1(e13, e13)))), introduced(avatar_definition, [new_symbols(naming, [spl68_388])])).
fof(f12831, plain, (~ spl68_125 | spl68_402 | ~ spl68_519), inference(avatar_contradiction_clause, [], [f12830])).
fof(f12830, plain, ($false | (~ spl68_125 | spl68_402 | ~ spl68_519)), inference(subsumption_resolution, [], [f12829, f3759])).
fof(f3759, plain, ((e24 = h1(e11)) | ~ spl68_519), inference(avatar_component_clause, [], [f3758])).
fof(f3758, plain, (spl68_519 <=> (e24 = h1(e11))), introduced(avatar_definition, [new_symbols(naming, [spl68_519])])).
fof(f12829, plain, (~ (e24 = h1(e11)) | (~ spl68_125 | spl68_402)), inference(forward_demodulation, [], [f12828, f758])).
fof(f12828, plain, (~ (h1(e11) = h5(e14)) | (~ spl68_125 | spl68_402)), inference(forward_demodulation, [], [f2985, f1411])).
fof(f1411, plain, ((op1(e10, e10) = e14) | ~ spl68_125), inference(avatar_component_clause, [], [f1409])).
fof(f1409, plain, (spl68_125 <=> (op1(e10, e10) = e14)), introduced(avatar_definition, [new_symbols(naming, [spl68_125])])).
fof(f2985, plain, (~ (h1(e11) = h5(op1(e10, e10))) | spl68_402), inference(avatar_component_clause, [], [f2983])).
fof(f2983, plain, (spl68_402 <=> (h1(e11) = h5(op1(e10, e10)))), introduced(avatar_definition, [new_symbols(naming, [spl68_402])])).
fof(f12817, plain, (~ spl68_107 | ~ spl68_237 | ~ spl68_314 | spl68_399), inference(avatar_contradiction_clause, [], [f12816])).
fof(f12816, plain, ($false | (~ spl68_107 | ~ spl68_237 | ~ spl68_314 | spl68_399)), inference(subsumption_resolution, [], [f12815, f1932])).
fof(f1932, plain, ((e21 = op2(e20, e23)) | ~ spl68_237), inference(avatar_component_clause, [], [f1930])).
fof(f1930, plain, (spl68_237 <=> (e21 = op2(e20, e23))), introduced(avatar_definition, [new_symbols(naming, [spl68_237])])).
fof(f12815, plain, (~ (e21 = op2(e20, e23)) | (~ spl68_107 | ~ spl68_314 | spl68_399)), inference(forward_demodulation, [], [f12814, f2525])).
fof(f2525, plain, ((e21 = h5(e11)) | ~ spl68_314), inference(avatar_component_clause, [], [f2524])).
fof(f2524, plain, (spl68_314 <=> (e21 = h5(e11))), introduced(avatar_definition, [new_symbols(naming, [spl68_314])])).
fof(f12814, plain, (~ (op2(e20, e23) = h5(e11)) | (~ spl68_107 | spl68_399)), inference(forward_demodulation, [], [f2973, f1336])).
fof(f1336, plain, ((e11 = op1(e10, e13)) | ~ spl68_107), inference(avatar_component_clause, [], [f1334])).
fof(f1334, plain, (spl68_107 <=> (e11 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl68_107])])).
fof(f2973, plain, (~ (op2(e20, e23) = h5(op1(e10, e13))) | spl68_399), inference(avatar_component_clause, [], [f2971])).
fof(f2971, plain, (spl68_399 <=> (op2(e20, e23) = h5(op1(e10, e13)))), introduced(avatar_definition, [new_symbols(naming, [spl68_399])])).
fof(f12627, plain, (~ spl68_240 | ~ spl68_432), inference(avatar_split_clause, [], [f12614, f3224, f1942])).
fof(f1942, plain, (spl68_240 <=> (e24 = op2(e20, e23))), introduced(avatar_definition, [new_symbols(naming, [spl68_240])])).
fof(f12614, plain, (~ (e24 = op2(e20, e23)) | ~ spl68_432), inference(backward_demodulation, [], [f2462, f3225])).
fof(f2462, plain, ~ (op2(e20, e23) = h4(e11)), inference(backward_demodulation, [], [f467, f755])).
fof(f755, plain, (op2(e23, e23) = h4(e11)), inference(cnf_transformation, [], [f19])).
fof(f19, plain, ((op2(e23, op2(e23, e23)) = h4(e13)) & (h4(e12) = op2(op2(e23, op2(e23, e23)), e23)) & (op2(e23, e23) = h4(e11)) & (h4(e10) = op2(op2(e23, e23), e23)) & (e23 = h4(e14))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG052+1.p', ax19)).
fof(f467, plain, ~ (op2(e20, e23) = op2(e23, e23)), inference(cnf_transformation, [], [f8])).
fof(f8, plain, (~ (op2(e24, e23) = op2(e24, e24)) & ~ (op2(e24, e22) = op2(e24, e24)) & ~ (op2(e24, e22) = op2(e24, e23)) & ~ (op2(e24, e21) = op2(e24, e24)) & ~ (op2(e24, e21) = op2(e24, e23)) & ~ (op2(e24, e21) = op2(e24, e22)) & ~ (op2(e24, e20) = op2(e24, e24)) & ~ (op2(e24, e20) = op2(e24, e23)) & ~ (op2(e24, e20) = op2(e24, e22)) & ~ (op2(e24, e20) = op2(e24, e21)) & ~ (op2(e23, e23) = op2(e23, e24)) & ~ (op2(e23, e22) = op2(e23, e24)) & ~ (op2(e23, e22) = op2(e23, e23)) & ~ (op2(e23, e21) = op2(e23, e24)) & ~ (op2(e23, e21) = op2(e23, e23)) & ~ (op2(e23, e21) = op2(e23, e22)) & ~ (op2(e23, e20) = op2(e23, e24)) & ~ (op2(e23, e20) = op2(e23, e23)) & ~ (op2(e23, e20) = op2(e23, e22)) & ~ (op2(e23, e20) = op2(e23, e21)) & ~ (op2(e22, e23) = op2(e22, e24)) & ~ (op2(e22, e22) = op2(e22, e24)) & ~ (op2(e22, e22) = op2(e22, e23)) & ~ (op2(e22, e21) = op2(e22, e24)) & ~ (op2(e22, e21) = op2(e22, e23)) & ~ (op2(e22, e21) = op2(e22, e22)) & ~ (op2(e22, e20) = op2(e22, e24)) & ~ (op2(e22, e20) = op2(e22, e23)) & ~ (op2(e22, e20) = op2(e22, e22)) & ~ (op2(e22, e20) = op2(e22, e21)) & ~ (op2(e21, e23) = op2(e21, e24)) & ~ (op2(e21, e22) = op2(e21, e24)) & ~ (op2(e21, e22) = op2(e21, e23)) & ~ (op2(e21, e21) = op2(e21, e24)) & ~ (op2(e21, e21) = op2(e21, e23)) & ~ (op2(e21, e21) = op2(e21, e22)) & ~ (op2(e21, e20) = op2(e21, e24)) & ~ (op2(e21, e20) = op2(e21, e23)) & ~ (op2(e21, e20) = op2(e21, e22)) & ~ (op2(e21, e20) = op2(e21, e21)) & ~ (op2(e20, e23) = op2(e20, e24)) & ~ (op2(e20, e22) = op2(e20, e24)) & ~ (op2(e20, e22) = op2(e20, e23)) & ~ (op2(e20, e21) = op2(e20, e24)) & ~ (op2(e20, e21) = op2(e20, e23)) & ~ (op2(e20, e21) = op2(e20, e22)) & ~ (op2(e20, e20) = op2(e20, e24)) & ~ (op2(e20, e20) = op2(e20, e23)) & ~ (op2(e20, e20) = op2(e20, e22)) & ~ (op2(e20, e20) = op2(e20, e21)) & ~ (op2(e23, e24) = op2(e24, e24)) & ~ (op2(e22, e24) = op2(e24, e24)) & ~ (op2(e22, e24) = op2(e23, e24)) & ~ (op2(e21, e24) = op2(e24, e24)) & ~ (op2(e21, e24) = op2(e23, e24)) & ~ (op2(e21, e24) = op2(e22, e24)) & ~ (op2(e20, e24) = op2(e24, e24)) & ~ (op2(e20, e24) = op2(e23, e24)) & ~ (op2(e20, e24) = op2(e22, e24)) & ~ (op2(e20, e24) = op2(e21, e24)) & ~ (op2(e23, e23) = op2(e24, e23)) & ~ (op2(e22, e23) = op2(e24, e23)) & ~ (op2(e22, e23) = op2(e23, e23)) & ~ (op2(e21, e23) = op2(e24, e23)) & ~ (op2(e21, e23) = op2(e23, e23)) & ~ (op2(e21, e23) = op2(e22, e23)) & ~ (op2(e20, e23) = op2(e24, e23)) & ~ (op2(e20, e23) = op2(e23, e23)) & ~ (op2(e20, e23) = op2(e22, e23)) & ~ (op2(e20, e23) = op2(e21, e23)) & ~ (op2(e23, e22) = op2(e24, e22)) & ~ (op2(e22, e22) = op2(e24, e22)) & ~ (op2(e22, e22) = op2(e23, e22)) & ~ (op2(e21, e22) = op2(e24, e22)) & ~ (op2(e21, e22) = op2(e23, e22)) & ~ (op2(e21, e22) = op2(e22, e22)) & ~ (op2(e20, e22) = op2(e24, e22)) & ~ (op2(e20, e22) = op2(e23, e22)) & ~ (op2(e20, e22) = op2(e22, e22)) & ~ (op2(e20, e22) = op2(e21, e22)) & ~ (op2(e23, e21) = op2(e24, e21)) & ~ (op2(e22, e21) = op2(e24, e21)) & ~ (op2(e22, e21) = op2(e23, e21)) & ~ (op2(e21, e21) = op2(e24, e21)) & ~ (op2(e21, e21) = op2(e23, e21)) & ~ (op2(e21, e21) = op2(e22, e21)) & ~ (op2(e20, e21) = op2(e24, e21)) & ~ (op2(e20, e21) = op2(e23, e21)) & ~ (op2(e20, e21) = op2(e22, e21)) & ~ (op2(e20, e21) = op2(e21, e21)) & ~ (op2(e23, e20) = op2(e24, e20)) & ~ (op2(e22, e20) = op2(e24, e20)) & ~ (op2(e22, e20) = op2(e23, e20)) & ~ (op2(e21, e20) = op2(e24, e20)) & ~ (op2(e21, e20) = op2(e23, e20)) & ~ (op2(e21, e20) = op2(e22, e20)) & ~ (op2(e20, e20) = op2(e24, e20)) & ~ (op2(e20, e20) = op2(e23, e20)) & ~ (op2(e20, e20) = op2(e22, e20)) & ~ (op2(e20, e20) = op2(e21, e20))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG052+1.p', ax8)).
fof(f12612, plain, (~ spl68_180 | ~ spl68_107 | ~ spl68_352 | ~ spl68_359 | spl68_466 | ~ spl68_490), inference(avatar_split_clause, [], [f12611, f3580, f3454, f2756, f2720, f1334, f1690])).
fof(f1690, plain, (spl68_180 <=> (e24 = op2(e23, e20))), introduced(avatar_definition, [new_symbols(naming, [spl68_180])])).
fof(f2720, plain, (spl68_352 <=> (e23 = h2(e10))), introduced(avatar_definition, [new_symbols(naming, [spl68_352])])).
fof(f2756, plain, (spl68_359 <=> (e20 = h2(e13))), introduced(avatar_definition, [new_symbols(naming, [spl68_359])])).
fof(f3454, plain, (spl68_466 <=> (h2(op1(e10, e13)) = op2(h2(e10), h2(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl68_466])])).
fof(f3580, plain, (spl68_490 <=> (e24 = h2(e11))), introduced(avatar_definition, [new_symbols(naming, [spl68_490])])).
fof(f12611, plain, (~ (e24 = op2(e23, e20)) | (~ spl68_107 | ~ spl68_352 | ~ spl68_359 | spl68_466 | ~ spl68_490)), inference(forward_demodulation, [], [f12610, f3581])).
fof(f3581, plain, ((e24 = h2(e11)) | ~ spl68_490), inference(avatar_component_clause, [], [f3580])).
fof(f12610, plain, (~ (op2(e23, e20) = h2(e11)) | (~ spl68_107 | ~ spl68_352 | ~ spl68_359 | spl68_466)), inference(forward_demodulation, [], [f12609, f1336])).
fof(f12609, plain, (~ (op2(e23, e20) = h2(op1(e10, e13))) | (~ spl68_352 | ~ spl68_359 | spl68_466)), inference(forward_demodulation, [], [f12608, f2721])).
fof(f2721, plain, ((e23 = h2(e10)) | ~ spl68_352), inference(avatar_component_clause, [], [f2720])).
fof(f12608, plain, (~ (h2(op1(e10, e13)) = op2(h2(e10), e20)) | (~ spl68_359 | spl68_466)), inference(forward_demodulation, [], [f3456, f2757])).
fof(f2757, plain, ((e20 = h2(e13)) | ~ spl68_359), inference(avatar_component_clause, [], [f2756])).
fof(f3456, plain, (~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | spl68_466), inference(avatar_component_clause, [], [f3454])).
fof(f12607, plain, (spl68_372 | ~ spl68_153 | ~ spl68_519), inference(avatar_split_clause, [], [f12606, f3758, f1577, f2821])).
fof(f2821, plain, (spl68_372 <=> (e22 = h1(e10))), introduced(avatar_definition, [new_symbols(naming, [spl68_372])])).
fof(f1577, plain, (spl68_153 <=> (e22 = op2(e24, e20))), introduced(avatar_definition, [new_symbols(naming, [spl68_153])])).
fof(f12606, plain, ((e22 = h1(e10)) | (~ spl68_153 | ~ spl68_519)), inference(forward_demodulation, [], [f12591, f1579])).
fof(f1579, plain, ((e22 = op2(e24, e20)) | ~ spl68_153), inference(avatar_component_clause, [], [f1577])).
fof(f12591, plain, ((op2(e24, e20) = h1(e10)) | ~ spl68_519), inference(backward_demodulation, [], [f2438, f3759])).
fof(f2438, plain, (h1(e10) = op2(h1(e11), e20)), inference(forward_demodulation, [], [f739, f740])).
fof(f740, plain, (op2(e20, e20) = h1(e11)), inference(cnf_transformation, [], [f16])).
fof(f16, plain, ((op2(e20, op2(e20, e20)) = h1(e13)) & (h1(e12) = op2(op2(e20, op2(e20, e20)), e20)) & (op2(e20, e20) = h1(e11)) & (h1(e10) = op2(op2(e20, e20), e20)) & (e20 = h1(e14))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG052+1.p', ax16)).
fof(f739, plain, (h1(e10) = op2(op2(e20, e20), e20)), inference(cnf_transformation, [], [f16])).
fof(f12605, plain, (spl68_364 | ~ spl68_234 | ~ spl68_519), inference(avatar_split_clause, [], [f12604, f3758, f1917, f2781])).
fof(f2781, plain, (spl68_364 <=> (e23 = h1(e13))), introduced(avatar_definition, [new_symbols(naming, [spl68_364])])).
fof(f1917, plain, (spl68_234 <=> (e23 = op2(e20, e24))), introduced(avatar_definition, [new_symbols(naming, [spl68_234])])).
fof(f12604, plain, ((e23 = h1(e13)) | (~ spl68_234 | ~ spl68_519)), inference(forward_demodulation, [], [f12590, f1919])).
fof(f1919, plain, ((e23 = op2(e20, e24)) | ~ spl68_234), inference(avatar_component_clause, [], [f1917])).
fof(f12590, plain, ((op2(e20, e24) = h1(e13)) | ~ spl68_519), inference(backward_demodulation, [], [f2437, f3759])).
fof(f2437, plain, (h1(e13) = op2(e20, h1(e11))), inference(backward_demodulation, [], [f742, f740])).
fof(f742, plain, (op2(e20, op2(e20, e20)) = h1(e13)), inference(cnf_transformation, [], [f16])).
fof(f12602, plain, (~ spl68_180 | ~ spl68_519), inference(avatar_split_clause, [], [f12588, f3758, f1690])).
fof(f12588, plain, (~ (e24 = op2(e23, e20)) | ~ spl68_519), inference(backward_demodulation, [], [f2431, f3759])).
fof(f2431, plain, ~ (op2(e23, e20) = h1(e11)), inference(backward_demodulation, [], [f437, f740])).
fof(f437, plain, ~ (op2(e20, e20) = op2(e23, e20)), inference(cnf_transformation, [], [f8])).
fof(f12535, plain, (~ spl68_177 | ~ spl68_47 | ~ spl68_314 | spl68_391), inference(avatar_split_clause, [], [f12534, f2939, f2524, f1082, f1678])).
fof(f1678, plain, (spl68_177 <=> (e21 = op2(e23, e20))), introduced(avatar_definition, [new_symbols(naming, [spl68_177])])).
fof(f1082, plain, (spl68_47 <=> (e11 = op1(e13, e10))), introduced(avatar_definition, [new_symbols(naming, [spl68_47])])).
fof(f2939, plain, (spl68_391 <=> (op2(e23, e20) = h5(op1(e13, e10)))), introduced(avatar_definition, [new_symbols(naming, [spl68_391])])).
fof(f12534, plain, (~ (e21 = op2(e23, e20)) | (~ spl68_47 | ~ spl68_314 | spl68_391)), inference(forward_demodulation, [], [f12533, f2525])).
fof(f12533, plain, (~ (op2(e23, e20) = h5(e11)) | (~ spl68_47 | spl68_391)), inference(forward_demodulation, [], [f2941, f1084])).
fof(f1084, plain, ((e11 = op1(e13, e10)) | ~ spl68_47), inference(avatar_component_clause, [], [f1082])).
fof(f2941, plain, (~ (op2(e23, e20) = h5(op1(e13, e10))) | spl68_391), inference(avatar_component_clause, [], [f2939])).
fof(f12513, plain, (~ spl68_35 | ~ spl68_252 | ~ spl68_359 | spl68_478), inference(avatar_contradiction_clause, [], [f12512])).
fof(f12512, plain, ($false | (~ spl68_35 | ~ spl68_252 | ~ spl68_359 | spl68_478)), inference(subsumption_resolution, [], [f12511, f1995])).
fof(f1995, plain, ((op2(e20, e20) = e21) | ~ spl68_252), inference(avatar_component_clause, [], [f1993])).
fof(f1993, plain, (spl68_252 <=> (op2(e20, e20) = e21)), introduced(avatar_definition, [new_symbols(naming, [spl68_252])])).
fof(f12511, plain, (~ (op2(e20, e20) = e21) | (~ spl68_35 | ~ spl68_359 | spl68_478)), inference(forward_demodulation, [], [f12510, f743])).
fof(f743, plain, (e21 = h2(e14)), inference(cnf_transformation, [], [f17])).
fof(f17, plain, ((op2(e21, op2(e21, e21)) = h2(e13)) & (h2(e12) = op2(op2(e21, op2(e21, e21)), e21)) & (op2(e21, e21) = h2(e11)) & (h2(e10) = op2(op2(e21, e21), e21)) & (e21 = h2(e14))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG052+1.p', ax17)).
fof(f12510, plain, (~ (op2(e20, e20) = h2(e14)) | (~ spl68_35 | ~ spl68_359 | spl68_478)), inference(forward_demodulation, [], [f12509, f1033])).
fof(f12509, plain, (~ (op2(e20, e20) = h2(op1(e13, e13))) | (~ spl68_359 | spl68_478)), inference(forward_demodulation, [], [f3504, f2757])).
fof(f3504, plain, (~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | spl68_478), inference(avatar_component_clause, [], [f3502])).
fof(f3502, plain, (spl68_478 <=> (h2(op1(e13, e13)) = op2(h2(e13), h2(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl68_478])])).
fof(f12373, plain, (~ spl68_47 | ~ spl68_240 | ~ spl68_352 | ~ spl68_359 | spl68_475 | ~ spl68_490), inference(avatar_contradiction_clause, [], [f12372])).
fof(f12372, plain, ($false | (~ spl68_47 | ~ spl68_240 | ~ spl68_352 | ~ spl68_359 | spl68_475 | ~ spl68_490)), inference(subsumption_resolution, [], [f12367, f3581])).
fof(f12367, plain, (~ (e24 = h2(e11)) | (~ spl68_47 | ~ spl68_240 | ~ spl68_352 | ~ spl68_359 | spl68_475)), inference(backward_demodulation, [], [f12106, f1084])).
fof(f12106, plain, (~ (e24 = h2(op1(e13, e10))) | (~ spl68_240 | ~ spl68_352 | ~ spl68_359 | spl68_475)), inference(backward_demodulation, [], [f11877, f1944])).
fof(f1944, plain, ((e24 = op2(e20, e23)) | ~ spl68_240), inference(avatar_component_clause, [], [f1942])).
fof(f11877, plain, (~ (op2(e20, e23) = h2(op1(e13, e10))) | (~ spl68_352 | ~ spl68_359 | spl68_475)), inference(forward_demodulation, [], [f11876, f2757])).
fof(f11876, plain, (~ (h2(op1(e13, e10)) = op2(h2(e13), e23)) | (~ spl68_352 | spl68_475)), inference(forward_demodulation, [], [f3492, f2721])).
fof(f3492, plain, (~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | spl68_475), inference(avatar_component_clause, [], [f3490])).
fof(f3490, plain, (spl68_475 <=> (h2(op1(e13, e10)) = op2(h2(e13), h2(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl68_475])])).
fof(f12369, plain, (~ spl68_32 | ~ spl68_47), inference(avatar_split_clause, [], [f12361, f1082, f1019])).
fof(f1019, plain, (spl68_32 <=> (e11 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl68_32])])).
fof(f12361, plain, (~ (e11 = op1(e13, e13)) | ~ spl68_47), inference(backward_demodulation, [], [f417, f1084])).
fof(f417, plain, ~ (op1(e13, e10) = op1(e13, e13)), inference(cnf_transformation, [], [f7])).
fof(f7, plain, (~ (op1(e14, e13) = op1(e14, e14)) & ~ (op1(e14, e12) = op1(e14, e14)) & ~ (op1(e14, e12) = op1(e14, e13)) & ~ (op1(e14, e11) = op1(e14, e14)) & ~ (op1(e14, e11) = op1(e14, e13)) & ~ (op1(e14, e11) = op1(e14, e12)) & ~ (op1(e14, e10) = op1(e14, e14)) & ~ (op1(e14, e10) = op1(e14, e13)) & ~ (op1(e14, e10) = op1(e14, e12)) & ~ (op1(e14, e10) = op1(e14, e11)) & ~ (op1(e13, e13) = op1(e13, e14)) & ~ (op1(e13, e12) = op1(e13, e14)) & ~ (op1(e13, e12) = op1(e13, e13)) & ~ (op1(e13, e11) = op1(e13, e14)) & ~ (op1(e13, e11) = op1(e13, e13)) & ~ (op1(e13, e11) = op1(e13, e12)) & ~ (op1(e13, e10) = op1(e13, e14)) & ~ (op1(e13, e10) = op1(e13, e13)) & ~ (op1(e13, e10) = op1(e13, e12)) & ~ (op1(e13, e10) = op1(e13, e11)) & ~ (op1(e12, e13) = op1(e12, e14)) & ~ (op1(e12, e12) = op1(e12, e14)) & ~ (op1(e12, e12) = op1(e12, e13)) & ~ (op1(e12, e11) = op1(e12, e14)) & ~ (op1(e12, e11) = op1(e12, e13)) & ~ (op1(e12, e11) = op1(e12, e12)) & ~ (op1(e12, e10) = op1(e12, e14)) & ~ (op1(e12, e10) = op1(e12, e13)) & ~ (op1(e12, e10) = op1(e12, e12)) & ~ (op1(e12, e10) = op1(e12, e11)) & ~ (op1(e11, e13) = op1(e11, e14)) & ~ (op1(e11, e12) = op1(e11, e14)) & ~ (op1(e11, e12) = op1(e11, e13)) & ~ (op1(e11, e11) = op1(e11, e14)) & ~ (op1(e11, e11) = op1(e11, e13)) & ~ (op1(e11, e11) = op1(e11, e12)) & ~ (op1(e11, e10) = op1(e11, e14)) & ~ (op1(e11, e10) = op1(e11, e13)) & ~ (op1(e11, e10) = op1(e11, e12)) & ~ (op1(e11, e10) = op1(e11, e11)) & ~ (op1(e10, e13) = op1(e10, e14)) & ~ (op1(e10, e12) = op1(e10, e14)) & ~ (op1(e10, e12) = op1(e10, e13)) & ~ (op1(e10, e11) = op1(e10, e14)) & ~ (op1(e10, e11) = op1(e10, e13)) & ~ (op1(e10, e11) = op1(e10, e12)) & ~ (op1(e10, e10) = op1(e10, e14)) & ~ (op1(e10, e10) = op1(e10, e13)) & ~ (op1(e10, e10) = op1(e10, e12)) & ~ (op1(e10, e10) = op1(e10, e11)) & ~ (op1(e13, e14) = op1(e14, e14)) & ~ (op1(e12, e14) = op1(e14, e14)) & ~ (op1(e12, e14) = op1(e13, e14)) & ~ (op1(e11, e14) = op1(e14, e14)) & ~ (op1(e11, e14) = op1(e13, e14)) & ~ (op1(e11, e14) = op1(e12, e14)) & ~ (op1(e10, e14) = op1(e14, e14)) & ~ (op1(e10, e14) = op1(e13, e14)) & ~ (op1(e10, e14) = op1(e12, e14)) & ~ (op1(e10, e14) = op1(e11, e14)) & ~ (op1(e13, e13) = op1(e14, e13)) & ~ (op1(e12, e13) = op1(e14, e13)) & ~ (op1(e12, e13) = op1(e13, e13)) & ~ (op1(e11, e13) = op1(e14, e13)) & ~ (op1(e11, e13) = op1(e13, e13)) & ~ (op1(e11, e13) = op1(e12, e13)) & ~ (op1(e10, e13) = op1(e14, e13)) & ~ (op1(e10, e13) = op1(e13, e13)) & ~ (op1(e10, e13) = op1(e12, e13)) & ~ (op1(e10, e13) = op1(e11, e13)) & ~ (op1(e13, e12) = op1(e14, e12)) & ~ (op1(e12, e12) = op1(e14, e12)) & ~ (op1(e12, e12) = op1(e13, e12)) & ~ (op1(e11, e12) = op1(e14, e12)) & ~ (op1(e11, e12) = op1(e13, e12)) & ~ (op1(e11, e12) = op1(e12, e12)) & ~ (op1(e10, e12) = op1(e14, e12)) & ~ (op1(e10, e12) = op1(e13, e12)) & ~ (op1(e10, e12) = op1(e12, e12)) & ~ (op1(e10, e12) = op1(e11, e12)) & ~ (op1(e13, e11) = op1(e14, e11)) & ~ (op1(e12, e11) = op1(e14, e11)) & ~ (op1(e12, e11) = op1(e13, e11)) & ~ (op1(e11, e11) = op1(e14, e11)) & ~ (op1(e11, e11) = op1(e13, e11)) & ~ (op1(e11, e11) = op1(e12, e11)) & ~ (op1(e10, e11) = op1(e14, e11)) & ~ (op1(e10, e11) = op1(e13, e11)) & ~ (op1(e10, e11) = op1(e12, e11)) & ~ (op1(e10, e11) = op1(e11, e11)) & ~ (op1(e13, e10) = op1(e14, e10)) & ~ (op1(e12, e10) = op1(e14, e10)) & ~ (op1(e12, e10) = op1(e13, e10)) & ~ (op1(e11, e10) = op1(e14, e10)) & ~ (op1(e11, e10) = op1(e13, e10)) & ~ (op1(e11, e10) = op1(e12, e10)) & ~ (op1(e10, e10) = op1(e14, e10)) & ~ (op1(e10, e10) = op1(e13, e10)) & ~ (op1(e10, e10) = op1(e12, e10)) & ~ (op1(e10, e10) = op1(e11, e10))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG052+1.p', ax7)).
fof(f12320, plain, (~ spl68_62 | ~ spl68_67), inference(avatar_split_clause, [], [f12309, f1166, f1145])).
fof(f1145, plain, (spl68_62 <=> (e11 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl68_62])])).
fof(f1166, plain, (spl68_67 <=> (e11 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl68_67])])).
fof(f12309, plain, (~ (e11 = op1(e12, e12)) | ~ spl68_67), inference(backward_demodulation, [], [f409, f1168])).
fof(f1168, plain, ((e11 = op1(e12, e11)) | ~ spl68_67), inference(avatar_component_clause, [], [f1166])).
fof(f409, plain, ~ (op1(e12, e11) = op1(e12, e12)), inference(cnf_transformation, [], [f7])).
fof(f12304, plain, (~ spl68_56 | ~ spl68_71), inference(avatar_split_clause, [], [f12294, f1183, f1120])).
fof(f1120, plain, (spl68_56 <=> (e10 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl68_56])])).
fof(f1183, plain, (spl68_71 <=> (e10 = op1(e12, e10))), introduced(avatar_definition, [new_symbols(naming, [spl68_71])])).
fof(f12294, plain, (~ (e10 = op1(e12, e13)) | ~ spl68_71), inference(backward_demodulation, [], [f407, f1185])).
fof(f1185, plain, ((e10 = op1(e12, e10)) | ~ spl68_71), inference(avatar_component_clause, [], [f1183])).
fof(f407, plain, ~ (op1(e12, e10) = op1(e12, e13)), inference(cnf_transformation, [], [f7])).
fof(f12288, plain, (~ spl68_8 | ~ spl68_83), inference(avatar_split_clause, [], [f12280, f1233, f918])).
fof(f918, plain, (spl68_8 <=> (e12 = op1(e14, e13))), introduced(avatar_definition, [new_symbols(naming, [spl68_8])])).
fof(f1233, plain, (spl68_83 <=> (e12 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl68_83])])).
fof(f12280, plain, (~ (e12 = op1(e14, e13)) | ~ spl68_83), inference(backward_demodulation, [], [f371, f1235])).
fof(f1235, plain, ((e12 = op1(e11, e13)) | ~ spl68_83), inference(avatar_component_clause, [], [f1233])).
fof(f371, plain, ~ (op1(e11, e13) = op1(e14, e13)), inference(cnf_transformation, [], [f7])).
fof(f12274, plain, (~ spl68_62 | ~ spl68_87), inference(avatar_split_clause, [], [f12262, f1250, f1145])).
fof(f1250, plain, (spl68_87 <=> (e11 = op1(e11, e12))), introduced(avatar_definition, [new_symbols(naming, [spl68_87])])).
fof(f12262, plain, (~ (e11 = op1(e12, e12)) | ~ spl68_87), inference(backward_demodulation, [], [f359, f1252])).
fof(f1252, plain, ((e11 = op1(e11, e12)) | ~ spl68_87), inference(avatar_component_clause, [], [f1250])).
fof(f359, plain, ~ (op1(e11, e12) = op1(e12, e12)), inference(cnf_transformation, [], [f7])).
fof(f12257, plain, (~ spl68_85 | ~ spl68_95), inference(avatar_split_clause, [], [f12250, f1283, f1241])).
fof(f1241, plain, (spl68_85 <=> (e14 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl68_85])])).
fof(f1283, plain, (spl68_95 <=> (e14 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl68_95])])).
fof(f12250, plain, (~ (e14 = op1(e11, e13)) | ~ spl68_95), inference(backward_demodulation, [], [f400, f1285])).
fof(f1285, plain, ((e14 = op1(e11, e11)) | ~ spl68_95), inference(avatar_component_clause, [], [f1283])).
fof(f400, plain, ~ (op1(e11, e11) = op1(e11, e13)), inference(cnf_transformation, [], [f7])).
fof(f12245, plain, (~ spl68_49 | ~ spl68_99), inference(avatar_split_clause, [], [f12233, f1300, f1090])).
fof(f1090, plain, (spl68_49 <=> (e13 = op1(e13, e10))), introduced(avatar_definition, [new_symbols(naming, [spl68_49])])).
fof(f1300, plain, (spl68_99 <=> (e13 = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl68_99])])).
fof(f12233, plain, (~ (e13 = op1(e13, e10)) | ~ spl68_99), inference(backward_demodulation, [], [f340, f1302])).
fof(f1302, plain, ((e13 = op1(e11, e10)) | ~ spl68_99), inference(avatar_component_clause, [], [f1300])).
fof(f340, plain, ~ (op1(e11, e10) = op1(e13, e10)), inference(cnf_transformation, [], [f7])).
fof(f12166, plain, (~ spl68_113 | ~ spl68_118), inference(avatar_split_clause, [], [f12155, f1380, f1359])).
fof(f1359, plain, (spl68_113 <=> (e12 = op1(e10, e12))), introduced(avatar_definition, [new_symbols(naming, [spl68_113])])).
fof(f1380, plain, (spl68_118 <=> (e12 = op1(e10, e11))), introduced(avatar_definition, [new_symbols(naming, [spl68_118])])).
fof(f12155, plain, (~ (e12 = op1(e10, e12)) | ~ spl68_118), inference(backward_demodulation, [], [f389, f1382])).
fof(f1382, plain, ((e12 = op1(e10, e11)) | ~ spl68_118), inference(avatar_component_clause, [], [f1380])).
fof(f389, plain, ~ (op1(e10, e11) = op1(e10, e12)), inference(cnf_transformation, [], [f7])).
fof(f12136, plain, (spl68_15 | ~ spl68_128), inference(avatar_split_clause, [], [f12124, f1422, f947])).
fof(f947, plain, (spl68_15 <=> (e14 = op1(e14, e12))), introduced(avatar_definition, [new_symbols(naming, [spl68_15])])).
fof(f1422, plain, (spl68_128 <=> (e12 = unit1)), introduced(avatar_definition, [new_symbols(naming, [spl68_128])])).
fof(f12124, plain, ((e14 = op1(e14, e12)) | ~ spl68_128), inference(backward_demodulation, [], [f197, f1424])).
fof(f1424, plain, ((e12 = unit1) | ~ spl68_128), inference(avatar_component_clause, [], [f1422])).
fof(f197, plain, (e14 = op1(e14, unit1)), inference(cnf_transformation, [], [f2])).
fof(f2, plain, (((e14 = unit1) | (e13 = unit1) | (e12 = unit1) | (e11 = unit1) | (e10 = unit1)) & (e14 = op1(e14, unit1)) & (e14 = op1(unit1, e14)) & (e13 = op1(e13, unit1)) & (e13 = op1(unit1, e13)) & (e12 = op1(e12, unit1)) & (e12 = op1(unit1, e12)) & (e11 = op1(e11, unit1)) & (e11 = op1(unit1, e11)) & (e10 = op1(e10, unit1)) & (e10 = op1(unit1, e10))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG052+1.p', ax2)).
fof(f12135, plain, (spl68_55 | ~ spl68_128), inference(avatar_split_clause, [], [f12123, f1422, f1115])).
fof(f1115, plain, (spl68_55 <=> (e14 = op1(e12, e14))), introduced(avatar_definition, [new_symbols(naming, [spl68_55])])).
fof(f12123, plain, ((e14 = op1(e12, e14)) | ~ spl68_128), inference(backward_demodulation, [], [f196, f1424])).
fof(f196, plain, (e14 = op1(unit1, e14)), inference(cnf_transformation, [], [f2])).
fof(f12134, plain, (spl68_39 | ~ spl68_128), inference(avatar_split_clause, [], [f12122, f1422, f1048])).
fof(f1048, plain, (spl68_39 <=> (e13 = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl68_39])])).
fof(f12122, plain, ((e13 = op1(e13, e12)) | ~ spl68_128), inference(backward_demodulation, [], [f195, f1424])).
fof(f195, plain, (e13 = op1(e13, unit1)), inference(cnf_transformation, [], [f2])).
fof(f12133, plain, (spl68_59 | ~ spl68_128), inference(avatar_split_clause, [], [f12121, f1422, f1132])).
fof(f1132, plain, (spl68_59 <=> (e13 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl68_59])])).
fof(f12121, plain, ((e13 = op1(e12, e13)) | ~ spl68_128), inference(backward_demodulation, [], [f194, f1424])).
fof(f194, plain, (e13 = op1(unit1, e13)), inference(cnf_transformation, [], [f2])).
fof(f12132, plain, (spl68_63 | ~ spl68_128), inference(avatar_split_clause, [], [f12120, f1422, f1149])).
fof(f1149, plain, (spl68_63 <=> (e12 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl68_63])])).
fof(f12120, plain, ((e12 = op1(e12, e12)) | ~ spl68_128), inference(backward_demodulation, [], [f193, f1424])).
fof(f193, plain, (e12 = op1(e12, unit1)), inference(cnf_transformation, [], [f2])).
fof(f12130, plain, (spl68_87 | ~ spl68_128), inference(avatar_split_clause, [], [f12118, f1422, f1250])).
fof(f12118, plain, ((e11 = op1(e11, e12)) | ~ spl68_128), inference(backward_demodulation, [], [f191, f1424])).
fof(f191, plain, (e11 = op1(e11, unit1)), inference(cnf_transformation, [], [f2])).
fof(f12129, plain, (spl68_67 | ~ spl68_128), inference(avatar_split_clause, [], [f12117, f1422, f1166])).
fof(f12117, plain, ((e11 = op1(e12, e11)) | ~ spl68_128), inference(backward_demodulation, [], [f190, f1424])).
fof(f190, plain, (e11 = op1(unit1, e11)), inference(cnf_transformation, [], [f2])).
fof(f12128, plain, (spl68_111 | ~ spl68_128), inference(avatar_split_clause, [], [f12116, f1422, f1351])).
fof(f1351, plain, (spl68_111 <=> (e10 = op1(e10, e12))), introduced(avatar_definition, [new_symbols(naming, [spl68_111])])).
fof(f12116, plain, ((e10 = op1(e10, e12)) | ~ spl68_128), inference(backward_demodulation, [], [f189, f1424])).
fof(f189, plain, (e10 = op1(e10, unit1)), inference(cnf_transformation, [], [f2])).
fof(f12127, plain, (spl68_71 | ~ spl68_128), inference(avatar_split_clause, [], [f12115, f1422, f1183])).
fof(f12115, plain, ((e10 = op1(e12, e10)) | ~ spl68_128), inference(backward_demodulation, [], [f188, f1424])).
fof(f188, plain, (e10 = op1(unit1, e10)), inference(cnf_transformation, [], [f2])).
fof(f12077, plain, (spl68_322 | ~ spl68_213 | ~ spl68_326), inference(avatar_split_clause, [], [f12076, f2589, f1829, f2569])).
fof(f2569, plain, (spl68_322 <=> (e22 = h4(e10))), introduced(avatar_definition, [new_symbols(naming, [spl68_322])])).
fof(f1829, plain, (spl68_213 <=> (e22 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl68_213])])).
fof(f2589, plain, (spl68_326 <=> (e21 = h4(e11))), introduced(avatar_definition, [new_symbols(naming, [spl68_326])])).
fof(f12076, plain, ((e22 = h4(e10)) | (~ spl68_213 | ~ spl68_326)), inference(forward_demodulation, [], [f12059, f1831])).
fof(f1831, plain, ((e22 = op2(e21, e23)) | ~ spl68_213), inference(avatar_component_clause, [], [f1829])).
fof(f12059, plain, ((op2(e21, e23) = h4(e10)) | ~ spl68_326), inference(backward_demodulation, [], [f2471, f2590])).
fof(f2590, plain, ((e21 = h4(e11)) | ~ spl68_326), inference(avatar_component_clause, [], [f2589])).
fof(f2471, plain, (h4(e10) = op2(h4(e11), e23)), inference(forward_demodulation, [], [f754, f755])).
fof(f754, plain, (h4(e10) = op2(op2(e23, e23), e23)), inference(cnf_transformation, [], [f19])).
fof(f12075, plain, (~ spl68_177 | ~ spl68_326), inference(avatar_split_clause, [], [f12058, f2589, f1678])).
fof(f12058, plain, (~ (e21 = op2(e23, e20)) | ~ spl68_326), inference(backward_demodulation, [], [f2466, f2590])).
fof(f2466, plain, ~ (op2(e23, e20) = h4(e11)), inference(backward_demodulation, [], [f517, f755])).
fof(f517, plain, ~ (op2(e23, e20) = op2(e23, e23)), inference(cnf_transformation, [], [f8])).
fof(f12017, plain, (spl68_367 | ~ spl68_229 | ~ spl68_376), inference(avatar_split_clause, [], [f12016, f2841, f1896, f2796])).
fof(f2796, plain, (spl68_367 <=> (e23 = h1(e10))), introduced(avatar_definition, [new_symbols(naming, [spl68_367])])).
fof(f1896, plain, (spl68_229 <=> (e23 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl68_229])])).
fof(f2841, plain, (spl68_376 <=> (e21 = h1(e11))), introduced(avatar_definition, [new_symbols(naming, [spl68_376])])).
fof(f12016, plain, ((e23 = h1(e10)) | (~ spl68_229 | ~ spl68_376)), inference(forward_demodulation, [], [f12001, f1898])).
fof(f1898, plain, ((e23 = op2(e21, e20)) | ~ spl68_229), inference(avatar_component_clause, [], [f1896])).
fof(f12001, plain, ((op2(e21, e20) = h1(e10)) | ~ spl68_376), inference(backward_demodulation, [], [f2438, f2842])).
fof(f2842, plain, ((e21 = h1(e11)) | ~ spl68_376), inference(avatar_component_clause, [], [f2841])).
fof(f12014, plain, (~ spl68_177 | ~ spl68_376), inference(avatar_split_clause, [], [f11999, f2841, f1678])).
fof(f11999, plain, (~ (e21 = op2(e23, e20)) | ~ spl68_376), inference(backward_demodulation, [], [f2431, f2842])).
fof(f11870, plain, (~ spl68_62 | ~ spl68_225 | ~ spl68_375 | spl68_502 | ~ spl68_519), inference(avatar_contradiction_clause, [], [f11869])).
fof(f11869, plain, ($false | (~ spl68_62 | ~ spl68_225 | ~ spl68_375 | spl68_502 | ~ spl68_519)), inference(subsumption_resolution, [], [f11868, f1881])).
fof(f1881, plain, ((e24 = op2(e21, e21)) | ~ spl68_225), inference(avatar_component_clause, [], [f1879])).
fof(f1879, plain, (spl68_225 <=> (e24 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl68_225])])).
fof(f11868, plain, (~ (e24 = op2(e21, e21)) | (~ spl68_62 | ~ spl68_375 | spl68_502 | ~ spl68_519)), inference(forward_demodulation, [], [f11867, f3759])).
fof(f11867, plain, (~ (op2(e21, e21) = h1(e11)) | (~ spl68_62 | ~ spl68_375 | spl68_502)), inference(forward_demodulation, [], [f11866, f1147])).
fof(f1147, plain, ((e11 = op1(e12, e12)) | ~ spl68_62), inference(avatar_component_clause, [], [f1145])).
fof(f11866, plain, (~ (op2(e21, e21) = h1(op1(e12, e12))) | (~ spl68_375 | spl68_502)), inference(forward_demodulation, [], [f3662, f2837])).
fof(f2837, plain, ((e21 = h1(e12)) | ~ spl68_375), inference(avatar_component_clause, [], [f2836])).
fof(f2836, plain, (spl68_375 <=> (e21 = h1(e12))), introduced(avatar_definition, [new_symbols(naming, [spl68_375])])).
fof(f3662, plain, (~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | spl68_502), inference(avatar_component_clause, [], [f3660])).
fof(f3660, plain, (spl68_502 <=> (h1(op1(e12, e12)) = op2(h1(e12), h1(e12)))), introduced(avatar_definition, [new_symbols(naming, [spl68_502])])).
fof(f11857, plain, (~ spl68_56 | ~ spl68_213 | ~ spl68_364 | ~ spl68_372 | ~ spl68_375 | spl68_503), inference(avatar_contradiction_clause, [], [f11856])).
fof(f11856, plain, ($false | (~ spl68_56 | ~ spl68_213 | ~ spl68_364 | ~ spl68_372 | ~ spl68_375 | spl68_503)), inference(subsumption_resolution, [], [f11855, f1831])).
fof(f11855, plain, (~ (e22 = op2(e21, e23)) | (~ spl68_56 | ~ spl68_364 | ~ spl68_372 | ~ spl68_375 | spl68_503)), inference(forward_demodulation, [], [f11854, f2822])).
fof(f2822, plain, ((e22 = h1(e10)) | ~ spl68_372), inference(avatar_component_clause, [], [f2821])).
fof(f11854, plain, (~ (op2(e21, e23) = h1(e10)) | (~ spl68_56 | ~ spl68_364 | ~ spl68_375 | spl68_503)), inference(forward_demodulation, [], [f11853, f1122])).
fof(f1122, plain, ((e10 = op1(e12, e13)) | ~ spl68_56), inference(avatar_component_clause, [], [f1120])).
fof(f11853, plain, (~ (op2(e21, e23) = h1(op1(e12, e13))) | (~ spl68_364 | ~ spl68_375 | spl68_503)), inference(forward_demodulation, [], [f11852, f2837])).
fof(f11852, plain, (~ (h1(op1(e12, e13)) = op2(h1(e12), e23)) | (~ spl68_364 | spl68_503)), inference(forward_demodulation, [], [f3666, f2782])).
fof(f2782, plain, ((e23 = h1(e13)) | ~ spl68_364), inference(avatar_component_clause, [], [f2781])).
fof(f3666, plain, (~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | spl68_503), inference(avatar_component_clause, [], [f3664])).
fof(f3664, plain, (spl68_503 <=> (h1(op1(e12, e13)) = op2(h1(e12), h1(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl68_503])])).
fof(f11851, plain, (~ spl68_70 | ~ spl68_206 | ~ spl68_375 | spl68_501 | ~ spl68_519), inference(avatar_contradiction_clause, [], [f11850])).
fof(f11850, plain, ($false | (~ spl68_70 | ~ spl68_206 | ~ spl68_375 | spl68_501 | ~ spl68_519)), inference(subsumption_resolution, [], [f11849, f1802])).
fof(f1802, plain, ((e20 = op2(e21, e24)) | ~ spl68_206), inference(avatar_component_clause, [], [f1800])).
fof(f1800, plain, (spl68_206 <=> (e20 = op2(e21, e24))), introduced(avatar_definition, [new_symbols(naming, [spl68_206])])).
fof(f11849, plain, (~ (e20 = op2(e21, e24)) | (~ spl68_70 | ~ spl68_375 | spl68_501 | ~ spl68_519)), inference(forward_demodulation, [], [f11848, f738])).
fof(f738, plain, (e20 = h1(e14)), inference(cnf_transformation, [], [f16])).
fof(f11848, plain, (~ (op2(e21, e24) = h1(e14)) | (~ spl68_70 | ~ spl68_375 | spl68_501 | ~ spl68_519)), inference(forward_demodulation, [], [f11847, f1180])).
fof(f1180, plain, ((e14 = op1(e12, e11)) | ~ spl68_70), inference(avatar_component_clause, [], [f1178])).
fof(f1178, plain, (spl68_70 <=> (e14 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl68_70])])).
fof(f11847, plain, (~ (op2(e21, e24) = h1(op1(e12, e11))) | (~ spl68_375 | spl68_501 | ~ spl68_519)), inference(forward_demodulation, [], [f11846, f2837])).
fof(f11846, plain, (~ (h1(op1(e12, e11)) = op2(h1(e12), e24)) | (spl68_501 | ~ spl68_519)), inference(forward_demodulation, [], [f3658, f3759])).
fof(f3658, plain, (~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | spl68_501), inference(avatar_component_clause, [], [f3656])).
fof(f3656, plain, (spl68_501 <=> (h1(op1(e12, e11)) = op2(h1(e12), h1(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl68_501])])).
fof(f11837, plain, (~ spl68_49 | ~ spl68_169 | ~ spl68_364 | ~ spl68_372 | spl68_504), inference(avatar_contradiction_clause, [], [f11836])).
fof(f11836, plain, ($false | (~ spl68_49 | ~ spl68_169 | ~ spl68_364 | ~ spl68_372 | spl68_504)), inference(subsumption_resolution, [], [f11835, f1646])).
fof(f1646, plain, ((e23 = op2(e23, e22)) | ~ spl68_169), inference(avatar_component_clause, [], [f1644])).
fof(f1644, plain, (spl68_169 <=> (e23 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl68_169])])).
fof(f11835, plain, (~ (e23 = op2(e23, e22)) | (~ spl68_49 | ~ spl68_364 | ~ spl68_372 | spl68_504)), inference(forward_demodulation, [], [f11834, f2782])).
fof(f11834, plain, (~ (op2(e23, e22) = h1(e13)) | (~ spl68_49 | ~ spl68_364 | ~ spl68_372 | spl68_504)), inference(forward_demodulation, [], [f11833, f1092])).
fof(f1092, plain, ((e13 = op1(e13, e10)) | ~ spl68_49), inference(avatar_component_clause, [], [f1090])).
fof(f11833, plain, (~ (op2(e23, e22) = h1(op1(e13, e10))) | (~ spl68_364 | ~ spl68_372 | spl68_504)), inference(forward_demodulation, [], [f11832, f2782])).
fof(f11832, plain, (~ (h1(op1(e13, e10)) = op2(h1(e13), e22)) | (~ spl68_372 | spl68_504)), inference(forward_demodulation, [], [f3670, f2822])).
fof(f3670, plain, (~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | spl68_504), inference(avatar_component_clause, [], [f3668])).
fof(f3668, plain, (spl68_504 <=> (h1(op1(e13, e10)) = op2(h1(e13), h1(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl68_504])])).
fof(f11831, plain, (~ spl68_73 | ~ spl68_217 | ~ spl68_372 | ~ spl68_375 | spl68_500), inference(avatar_contradiction_clause, [], [f11830])).
fof(f11830, plain, ($false | (~ spl68_73 | ~ spl68_217 | ~ spl68_372 | ~ spl68_375 | spl68_500)), inference(subsumption_resolution, [], [f11829, f1848])).
fof(f1848, plain, ((e21 = op2(e21, e22)) | ~ spl68_217), inference(avatar_component_clause, [], [f1846])).
fof(f1846, plain, (spl68_217 <=> (e21 = op2(e21, e22))), introduced(avatar_definition, [new_symbols(naming, [spl68_217])])).
fof(f11829, plain, (~ (e21 = op2(e21, e22)) | (~ spl68_73 | ~ spl68_372 | ~ spl68_375 | spl68_500)), inference(forward_demodulation, [], [f11828, f2837])).
fof(f11828, plain, (~ (op2(e21, e22) = h1(e12)) | (~ spl68_73 | ~ spl68_372 | ~ spl68_375 | spl68_500)), inference(forward_demodulation, [], [f11827, f1193])).
fof(f1193, plain, ((e12 = op1(e12, e10)) | ~ spl68_73), inference(avatar_component_clause, [], [f1191])).
fof(f1191, plain, (spl68_73 <=> (e12 = op1(e12, e10))), introduced(avatar_definition, [new_symbols(naming, [spl68_73])])).
fof(f11827, plain, (~ (op2(e21, e22) = h1(op1(e12, e10))) | (~ spl68_372 | ~ spl68_375 | spl68_500)), inference(forward_demodulation, [], [f11826, f2837])).
fof(f11826, plain, (~ (h1(op1(e12, e10)) = op2(h1(e12), e22)) | (~ spl68_372 | spl68_500)), inference(forward_demodulation, [], [f3654, f2822])).
fof(f3654, plain, (~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | spl68_500), inference(avatar_component_clause, [], [f3652])).
fof(f3652, plain, (spl68_500 <=> (h1(op1(e12, e10)) = op2(h1(e12), h1(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl68_500])])).
fof(f11817, plain, (~ spl68_41 | ~ spl68_158 | ~ spl68_364 | ~ spl68_372 | spl68_505 | ~ spl68_519), inference(avatar_contradiction_clause, [], [f11816])).
fof(f11816, plain, ($false | (~ spl68_41 | ~ spl68_158 | ~ spl68_364 | ~ spl68_372 | spl68_505 | ~ spl68_519)), inference(subsumption_resolution, [], [f11815, f1600])).
fof(f1600, plain, ((e22 = op2(e23, e24)) | ~ spl68_158), inference(avatar_component_clause, [], [f1598])).
fof(f1598, plain, (spl68_158 <=> (e22 = op2(e23, e24))), introduced(avatar_definition, [new_symbols(naming, [spl68_158])])).
fof(f11815, plain, (~ (e22 = op2(e23, e24)) | (~ spl68_41 | ~ spl68_364 | ~ spl68_372 | spl68_505 | ~ spl68_519)), inference(forward_demodulation, [], [f11814, f2822])).
fof(f11814, plain, (~ (op2(e23, e24) = h1(e10)) | (~ spl68_41 | ~ spl68_364 | spl68_505 | ~ spl68_519)), inference(forward_demodulation, [], [f11813, f1059])).
fof(f1059, plain, ((e10 = op1(e13, e11)) | ~ spl68_41), inference(avatar_component_clause, [], [f1057])).
fof(f1057, plain, (spl68_41 <=> (e10 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl68_41])])).
fof(f11813, plain, (~ (op2(e23, e24) = h1(op1(e13, e11))) | (~ spl68_364 | spl68_505 | ~ spl68_519)), inference(forward_demodulation, [], [f11812, f2782])).
fof(f11812, plain, (~ (h1(op1(e13, e11)) = op2(h1(e13), e24)) | (spl68_505 | ~ spl68_519)), inference(forward_demodulation, [], [f3674, f3759])).
fof(f3674, plain, (~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | spl68_505), inference(avatar_component_clause, [], [f3672])).
fof(f3672, plain, (spl68_505 <=> (h1(op1(e13, e11)) = op2(h1(e13), h1(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl68_505])])).
fof(f11811, plain, (~ spl68_85 | ~ spl68_136 | ~ spl68_364 | spl68_499 | ~ spl68_519), inference(avatar_contradiction_clause, [], [f11810])).
fof(f11810, plain, ($false | (~ spl68_85 | ~ spl68_136 | ~ spl68_364 | spl68_499 | ~ spl68_519)), inference(subsumption_resolution, [], [f11809, f1508])).
fof(f1508, plain, ((e20 = op2(e24, e23)) | ~ spl68_136), inference(avatar_component_clause, [], [f1506])).
fof(f1506, plain, (spl68_136 <=> (e20 = op2(e24, e23))), introduced(avatar_definition, [new_symbols(naming, [spl68_136])])).
fof(f11809, plain, (~ (e20 = op2(e24, e23)) | (~ spl68_85 | ~ spl68_364 | spl68_499 | ~ spl68_519)), inference(forward_demodulation, [], [f11808, f738])).
fof(f11808, plain, (~ (op2(e24, e23) = h1(e14)) | (~ spl68_85 | ~ spl68_364 | spl68_499 | ~ spl68_519)), inference(forward_demodulation, [], [f11807, f1243])).
fof(f1243, plain, ((e14 = op1(e11, e13)) | ~ spl68_85), inference(avatar_component_clause, [], [f1241])).
fof(f11807, plain, (~ (op2(e24, e23) = h1(op1(e11, e13))) | (~ spl68_364 | spl68_499 | ~ spl68_519)), inference(forward_demodulation, [], [f11806, f3759])).
fof(f11806, plain, (~ (h1(op1(e11, e13)) = op2(h1(e11), e23)) | (~ spl68_364 | spl68_499)), inference(forward_demodulation, [], [f3650, f2782])).
fof(f3650, plain, (~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | spl68_499), inference(avatar_component_clause, [], [f3648])).
fof(f3648, plain, (spl68_499 <=> (h1(op1(e11, e13)) = op2(h1(e11), h1(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl68_499])])).
fof(f11797, plain, (~ spl68_40 | ~ spl68_171 | ~ spl68_364 | ~ spl68_375 | spl68_506), inference(avatar_contradiction_clause, [], [f11796])).
fof(f11796, plain, ($false | (~ spl68_40 | ~ spl68_171 | ~ spl68_364 | ~ spl68_375 | spl68_506)), inference(subsumption_resolution, [], [f11795, f1655])).
fof(f1655, plain, ((e20 = op2(e23, e21)) | ~ spl68_171), inference(avatar_component_clause, [], [f1653])).
fof(f1653, plain, (spl68_171 <=> (e20 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl68_171])])).
fof(f11795, plain, (~ (e20 = op2(e23, e21)) | (~ spl68_40 | ~ spl68_364 | ~ spl68_375 | spl68_506)), inference(forward_demodulation, [], [f11794, f738])).
fof(f11794, plain, (~ (op2(e23, e21) = h1(e14)) | (~ spl68_40 | ~ spl68_364 | ~ spl68_375 | spl68_506)), inference(forward_demodulation, [], [f11793, f1054])).
fof(f1054, plain, ((e14 = op1(e13, e12)) | ~ spl68_40), inference(avatar_component_clause, [], [f1052])).
fof(f1052, plain, (spl68_40 <=> (e14 = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl68_40])])).
fof(f11793, plain, (~ (op2(e23, e21) = h1(op1(e13, e12))) | (~ spl68_364 | ~ spl68_375 | spl68_506)), inference(forward_demodulation, [], [f11792, f2782])).
fof(f11792, plain, (~ (h1(op1(e13, e12)) = op2(h1(e13), e21)) | (~ spl68_375 | spl68_506)), inference(forward_demodulation, [], [f3678, f2837])).
fof(f3678, plain, (~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | spl68_506), inference(avatar_component_clause, [], [f3676])).
fof(f3676, plain, (spl68_506 <=> (h1(op1(e13, e12)) = op2(h1(e13), h1(e12)))), introduced(avatar_definition, [new_symbols(naming, [spl68_506])])).
fof(f11791, plain, (~ spl68_89 | ~ spl68_149 | ~ spl68_364 | ~ spl68_375 | spl68_498 | ~ spl68_519), inference(avatar_contradiction_clause, [], [f11790])).
fof(f11790, plain, ($false | (~ spl68_89 | ~ spl68_149 | ~ spl68_364 | ~ spl68_375 | spl68_498 | ~ spl68_519)), inference(subsumption_resolution, [], [f11789, f1562])).
fof(f1562, plain, ((e23 = op2(e24, e21)) | ~ spl68_149), inference(avatar_component_clause, [], [f1560])).
fof(f1560, plain, (spl68_149 <=> (e23 = op2(e24, e21))), introduced(avatar_definition, [new_symbols(naming, [spl68_149])])).
fof(f11789, plain, (~ (e23 = op2(e24, e21)) | (~ spl68_89 | ~ spl68_364 | ~ spl68_375 | spl68_498 | ~ spl68_519)), inference(forward_demodulation, [], [f11788, f2782])).
fof(f11788, plain, (~ (op2(e24, e21) = h1(e13)) | (~ spl68_89 | ~ spl68_375 | spl68_498 | ~ spl68_519)), inference(forward_demodulation, [], [f11787, f1260])).
fof(f1260, plain, ((e13 = op1(e11, e12)) | ~ spl68_89), inference(avatar_component_clause, [], [f1258])).
fof(f1258, plain, (spl68_89 <=> (e13 = op1(e11, e12))), introduced(avatar_definition, [new_symbols(naming, [spl68_89])])).
fof(f11787, plain, (~ (op2(e24, e21) = h1(op1(e11, e12))) | (~ spl68_375 | spl68_498 | ~ spl68_519)), inference(forward_demodulation, [], [f11786, f3759])).
fof(f11786, plain, (~ (h1(op1(e11, e12)) = op2(h1(e11), e21)) | (~ spl68_375 | spl68_498)), inference(forward_demodulation, [], [f3646, f2837])).
fof(f3646, plain, (~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | spl68_498), inference(avatar_component_clause, [], [f3644])).
fof(f3644, plain, (spl68_498 <=> (h1(op1(e11, e12)) = op2(h1(e11), h1(e12)))), introduced(avatar_definition, [new_symbols(naming, [spl68_498])])).
fof(f11779, plain, (~ spl68_32 | ~ spl68_165 | ~ spl68_364 | spl68_507 | ~ spl68_519), inference(avatar_contradiction_clause, [], [f11778])).
fof(f11778, plain, ($false | (~ spl68_32 | ~ spl68_165 | ~ spl68_364 | spl68_507 | ~ spl68_519)), inference(subsumption_resolution, [], [f11777, f1629])).
fof(f1629, plain, ((e24 = op2(e23, e23)) | ~ spl68_165), inference(avatar_component_clause, [], [f1627])).
fof(f1627, plain, (spl68_165 <=> (e24 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl68_165])])).
fof(f11777, plain, (~ (e24 = op2(e23, e23)) | (~ spl68_32 | ~ spl68_364 | spl68_507 | ~ spl68_519)), inference(forward_demodulation, [], [f11776, f3759])).
fof(f11776, plain, (~ (op2(e23, e23) = h1(e11)) | (~ spl68_32 | ~ spl68_364 | spl68_507)), inference(forward_demodulation, [], [f11775, f1021])).
fof(f1021, plain, ((e11 = op1(e13, e13)) | ~ spl68_32), inference(avatar_component_clause, [], [f1019])).
fof(f11775, plain, (~ (op2(e23, e23) = h1(op1(e13, e13))) | (~ spl68_364 | spl68_507)), inference(forward_demodulation, [], [f3682, f2782])).
fof(f3682, plain, (~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | spl68_507), inference(avatar_component_clause, [], [f3680])).
fof(f3680, plain, (spl68_507 <=> (h1(op1(e13, e13)) = op2(h1(e13), h1(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl68_507])])).
fof(f11774, plain, (~ spl68_93 | ~ spl68_132 | ~ spl68_375 | spl68_497 | ~ spl68_519), inference(avatar_contradiction_clause, [], [f11773])).
fof(f11773, plain, ($false | (~ spl68_93 | ~ spl68_132 | ~ spl68_375 | spl68_497 | ~ spl68_519)), inference(subsumption_resolution, [], [f11772, f1491])).
fof(f1491, plain, ((e21 = op2(e24, e24)) | ~ spl68_132), inference(avatar_component_clause, [], [f1489])).
fof(f1489, plain, (spl68_132 <=> (e21 = op2(e24, e24))), introduced(avatar_definition, [new_symbols(naming, [spl68_132])])).
fof(f11772, plain, (~ (e21 = op2(e24, e24)) | (~ spl68_93 | ~ spl68_375 | spl68_497 | ~ spl68_519)), inference(forward_demodulation, [], [f11771, f2837])).
fof(f11771, plain, (~ (op2(e24, e24) = h1(e12)) | (~ spl68_93 | spl68_497 | ~ spl68_519)), inference(forward_demodulation, [], [f11770, f1277])).
fof(f1277, plain, ((e12 = op1(e11, e11)) | ~ spl68_93), inference(avatar_component_clause, [], [f1275])).
fof(f1275, plain, (spl68_93 <=> (e12 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl68_93])])).
fof(f11770, plain, (~ (op2(e24, e24) = h1(op1(e11, e11))) | (spl68_497 | ~ spl68_519)), inference(forward_demodulation, [], [f3642, f3759])).
fof(f3642, plain, (~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | spl68_497), inference(avatar_component_clause, [], [f3640])).
fof(f3640, plain, (spl68_497 <=> (h1(op1(e11, e11)) = op2(h1(e11), h1(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl68_497])])).
fof(f11765, plain, (~ spl68_2 | spl68_509), inference(avatar_contradiction_clause, [], [f11764])).
fof(f11764, plain, ($false | (~ spl68_2 | spl68_509)), inference(trivial_inequality_removal, [], [f11763])).
fof(f11763, plain, (~ (h1(e11) = h1(e11)) | (~ spl68_2 | spl68_509)), inference(forward_demodulation, [], [f3690, f895])).
fof(f895, plain, ((e11 = op1(e14, e14)) | ~ spl68_2), inference(avatar_component_clause, [], [f893])).
fof(f893, plain, (spl68_2 <=> (e11 = op1(e14, e14))), introduced(avatar_definition, [new_symbols(naming, [spl68_2])])).
fof(f3690, plain, (~ (h1(e11) = h1(op1(e14, e14))) | spl68_509), inference(avatar_component_clause, [], [f3688])).
fof(f3688, plain, (spl68_509 <=> (h1(e11) = h1(op1(e14, e14)))), introduced(avatar_definition, [new_symbols(naming, [spl68_509])])).
fof(f11762, plain, (~ spl68_97 | ~ spl68_145 | ~ spl68_372 | spl68_496 | ~ spl68_519), inference(avatar_contradiction_clause, [], [f11761])).
fof(f11761, plain, ($false | (~ spl68_97 | ~ spl68_145 | ~ spl68_372 | spl68_496 | ~ spl68_519)), inference(subsumption_resolution, [], [f11760, f1545])).
fof(f1545, plain, ((e24 = op2(e24, e22)) | ~ spl68_145), inference(avatar_component_clause, [], [f1543])).
fof(f1543, plain, (spl68_145 <=> (e24 = op2(e24, e22))), introduced(avatar_definition, [new_symbols(naming, [spl68_145])])).
fof(f11760, plain, (~ (e24 = op2(e24, e22)) | (~ spl68_97 | ~ spl68_372 | spl68_496 | ~ spl68_519)), inference(forward_demodulation, [], [f11759, f3759])).
fof(f11759, plain, (~ (op2(e24, e22) = h1(e11)) | (~ spl68_97 | ~ spl68_372 | spl68_496 | ~ spl68_519)), inference(forward_demodulation, [], [f11758, f1294])).
fof(f1294, plain, ((e11 = op1(e11, e10)) | ~ spl68_97), inference(avatar_component_clause, [], [f1292])).
fof(f1292, plain, (spl68_97 <=> (e11 = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl68_97])])).
fof(f11758, plain, (~ (op2(e24, e22) = h1(op1(e11, e10))) | (~ spl68_372 | spl68_496 | ~ spl68_519)), inference(forward_demodulation, [], [f11757, f3759])).
fof(f11757, plain, (~ (h1(op1(e11, e10)) = op2(h1(e11), e22)) | (~ spl68_372 | spl68_496)), inference(forward_demodulation, [], [f3638, f2822])).
fof(f3638, plain, (~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | spl68_496), inference(avatar_component_clause, [], [f3636])).
fof(f3636, plain, (spl68_496 <=> (h1(op1(e11, e10)) = op2(h1(e11), h1(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl68_496])])).
fof(f11749, plain, (~ spl68_8 | ~ spl68_237 | ~ spl68_364 | ~ spl68_375 | spl68_510), inference(avatar_contradiction_clause, [], [f11748])).
fof(f11748, plain, ($false | (~ spl68_8 | ~ spl68_237 | ~ spl68_364 | ~ spl68_375 | spl68_510)), inference(subsumption_resolution, [], [f11747, f1932])).
fof(f11747, plain, (~ (e21 = op2(e20, e23)) | (~ spl68_8 | ~ spl68_364 | ~ spl68_375 | spl68_510)), inference(forward_demodulation, [], [f11746, f2837])).
fof(f11746, plain, (~ (op2(e20, e23) = h1(e12)) | (~ spl68_8 | ~ spl68_364 | spl68_510)), inference(forward_demodulation, [], [f11745, f920])).
fof(f920, plain, ((e12 = op1(e14, e13)) | ~ spl68_8), inference(avatar_component_clause, [], [f918])).
fof(f11745, plain, (~ (op2(e20, e23) = h1(op1(e14, e13))) | (~ spl68_364 | spl68_510)), inference(forward_demodulation, [], [f3694, f2782])).
fof(f3694, plain, (~ (h1(op1(e14, e13)) = op2(e20, h1(e13))) | spl68_510), inference(avatar_component_clause, [], [f3692])).
fof(f3692, plain, (spl68_510 <=> (h1(op1(e14, e13)) = op2(e20, h1(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl68_510])])).
fof(f11744, plain, (~ spl68_109 | ~ spl68_189 | ~ spl68_364 | ~ spl68_372 | spl68_495), inference(avatar_contradiction_clause, [], [f11743])).
fof(f11743, plain, ($false | (~ spl68_109 | ~ spl68_189 | ~ spl68_364 | ~ spl68_372 | spl68_495)), inference(subsumption_resolution, [], [f11742, f1730])).
fof(f1730, plain, ((e23 = op2(e22, e23)) | ~ spl68_189), inference(avatar_component_clause, [], [f1728])).
fof(f1728, plain, (spl68_189 <=> (e23 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl68_189])])).
fof(f11742, plain, (~ (e23 = op2(e22, e23)) | (~ spl68_109 | ~ spl68_364 | ~ spl68_372 | spl68_495)), inference(forward_demodulation, [], [f11741, f2782])).
fof(f11741, plain, (~ (op2(e22, e23) = h1(e13)) | (~ spl68_109 | ~ spl68_364 | ~ spl68_372 | spl68_495)), inference(forward_demodulation, [], [f11740, f1344])).
fof(f1344, plain, ((e13 = op1(e10, e13)) | ~ spl68_109), inference(avatar_component_clause, [], [f1342])).
fof(f1342, plain, (spl68_109 <=> (e13 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl68_109])])).
fof(f11740, plain, (~ (op2(e22, e23) = h1(op1(e10, e13))) | (~ spl68_364 | ~ spl68_372 | spl68_495)), inference(forward_demodulation, [], [f11739, f2822])).
fof(f11739, plain, (~ (h1(op1(e10, e13)) = op2(h1(e10), e23)) | (~ spl68_364 | spl68_495)), inference(forward_demodulation, [], [f3634, f2782])).
fof(f3634, plain, (~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | spl68_495), inference(avatar_component_clause, [], [f3632])).
fof(f3632, plain, (spl68_495 <=> (h1(op1(e10, e13)) = op2(h1(e10), h1(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl68_495])])).
fof(f11731, plain, (~ spl68_11 | ~ spl68_248 | ~ spl68_372 | ~ spl68_375 | spl68_511), inference(avatar_contradiction_clause, [], [f11730])).
fof(f11730, plain, ($false | (~ spl68_11 | ~ spl68_248 | ~ spl68_372 | ~ spl68_375 | spl68_511)), inference(subsumption_resolution, [], [f11729, f1978])).
fof(f1978, plain, ((e22 = op2(e20, e21)) | ~ spl68_248), inference(avatar_component_clause, [], [f1976])).
fof(f1976, plain, (spl68_248 <=> (e22 = op2(e20, e21))), introduced(avatar_definition, [new_symbols(naming, [spl68_248])])).
fof(f11729, plain, (~ (e22 = op2(e20, e21)) | (~ spl68_11 | ~ spl68_372 | ~ spl68_375 | spl68_511)), inference(forward_demodulation, [], [f11728, f2822])).
fof(f11728, plain, (~ (op2(e20, e21) = h1(e10)) | (~ spl68_11 | ~ spl68_375 | spl68_511)), inference(forward_demodulation, [], [f11727, f933])).
fof(f933, plain, ((e10 = op1(e14, e12)) | ~ spl68_11), inference(avatar_component_clause, [], [f931])).
fof(f931, plain, (spl68_11 <=> (e10 = op1(e14, e12))), introduced(avatar_definition, [new_symbols(naming, [spl68_11])])).
fof(f11727, plain, (~ (op2(e20, e21) = h1(op1(e14, e12))) | (~ spl68_375 | spl68_511)), inference(forward_demodulation, [], [f3698, f2837])).
fof(f3698, plain, (~ (h1(op1(e14, e12)) = op2(e20, h1(e12))) | spl68_511), inference(avatar_component_clause, [], [f3696])).
fof(f3696, plain, (spl68_511 <=> (h1(op1(e14, e12)) = op2(e20, h1(e12)))), introduced(avatar_definition, [new_symbols(naming, [spl68_511])])).
fof(f11726, plain, (~ spl68_113 | ~ spl68_197 | ~ spl68_372 | ~ spl68_375 | spl68_494), inference(avatar_contradiction_clause, [], [f11725])).
fof(f11725, plain, ($false | (~ spl68_113 | ~ spl68_197 | ~ spl68_372 | ~ spl68_375 | spl68_494)), inference(subsumption_resolution, [], [f11724, f1764])).
fof(f1764, plain, ((e21 = op2(e22, e21)) | ~ spl68_197), inference(avatar_component_clause, [], [f1762])).
fof(f1762, plain, (spl68_197 <=> (e21 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl68_197])])).
fof(f11724, plain, (~ (e21 = op2(e22, e21)) | (~ spl68_113 | ~ spl68_372 | ~ spl68_375 | spl68_494)), inference(forward_demodulation, [], [f11723, f2837])).
fof(f11723, plain, (~ (op2(e22, e21) = h1(e12)) | (~ spl68_113 | ~ spl68_372 | ~ spl68_375 | spl68_494)), inference(forward_demodulation, [], [f11722, f1361])).
fof(f1361, plain, ((e12 = op1(e10, e12)) | ~ spl68_113), inference(avatar_component_clause, [], [f1359])).
fof(f11722, plain, (~ (op2(e22, e21) = h1(op1(e10, e12))) | (~ spl68_372 | ~ spl68_375 | spl68_494)), inference(forward_demodulation, [], [f11721, f2822])).
fof(f11721, plain, (~ (h1(op1(e10, e12)) = op2(h1(e10), e21)) | (~ spl68_375 | spl68_494)), inference(forward_demodulation, [], [f3630, f2837])).
fof(f3630, plain, (~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | spl68_494), inference(avatar_component_clause, [], [f3628])).
fof(f3628, plain, (spl68_494 <=> (h1(op1(e10, e12)) = op2(h1(e10), h1(e12)))), introduced(avatar_definition, [new_symbols(naming, [spl68_494])])).
fof(f11713, plain, (~ spl68_25 | ~ spl68_241 | ~ spl68_372 | spl68_513), inference(avatar_contradiction_clause, [], [f11712])).
fof(f11712, plain, ($false | (~ spl68_25 | ~ spl68_241 | ~ spl68_372 | spl68_513)), inference(subsumption_resolution, [], [f11711, f1949])).
fof(f1949, plain, ((e20 = op2(e20, e22)) | ~ spl68_241), inference(avatar_component_clause, [], [f1947])).
fof(f1947, plain, (spl68_241 <=> (e20 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl68_241])])).
fof(f11711, plain, (~ (e20 = op2(e20, e22)) | (~ spl68_25 | ~ spl68_372 | spl68_513)), inference(forward_demodulation, [], [f11710, f738])).
fof(f11710, plain, (~ (op2(e20, e22) = h1(e14)) | (~ spl68_25 | ~ spl68_372 | spl68_513)), inference(forward_demodulation, [], [f11709, f991])).
fof(f991, plain, ((e14 = op1(e14, e10)) | ~ spl68_25), inference(avatar_component_clause, [], [f989])).
fof(f989, plain, (spl68_25 <=> (e14 = op1(e14, e10))), introduced(avatar_definition, [new_symbols(naming, [spl68_25])])).
fof(f11709, plain, (~ (op2(e20, e22) = h1(op1(e14, e10))) | (~ spl68_372 | spl68_513)), inference(forward_demodulation, [], [f3706, f2822])).
fof(f3706, plain, (~ (h1(op1(e14, e10)) = op2(e20, h1(e10))) | spl68_513), inference(avatar_component_clause, [], [f3704])).
fof(f3704, plain, (spl68_513 <=> (h1(op1(e14, e10)) = op2(e20, h1(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl68_513])])).
fof(f11708, plain, (~ spl68_117 | ~ spl68_185 | ~ spl68_372 | spl68_493 | ~ spl68_519), inference(avatar_contradiction_clause, [], [f11707])).
fof(f11707, plain, ($false | (~ spl68_117 | ~ spl68_185 | ~ spl68_372 | spl68_493 | ~ spl68_519)), inference(subsumption_resolution, [], [f11706, f1713])).
fof(f1713, plain, ((e24 = op2(e22, e24)) | ~ spl68_185), inference(avatar_component_clause, [], [f1711])).
fof(f1711, plain, (spl68_185 <=> (e24 = op2(e22, e24))), introduced(avatar_definition, [new_symbols(naming, [spl68_185])])).
fof(f11706, plain, (~ (e24 = op2(e22, e24)) | (~ spl68_117 | ~ spl68_372 | spl68_493 | ~ spl68_519)), inference(forward_demodulation, [], [f11705, f3759])).
fof(f11705, plain, (~ (op2(e22, e24) = h1(e11)) | (~ spl68_117 | ~ spl68_372 | spl68_493 | ~ spl68_519)), inference(forward_demodulation, [], [f11704, f1378])).
fof(f1378, plain, ((e11 = op1(e10, e11)) | ~ spl68_117), inference(avatar_component_clause, [], [f1376])).
fof(f1376, plain, (spl68_117 <=> (e11 = op1(e10, e11))), introduced(avatar_definition, [new_symbols(naming, [spl68_117])])).
fof(f11704, plain, (~ (op2(e22, e24) = h1(op1(e10, e11))) | (~ spl68_372 | spl68_493 | ~ spl68_519)), inference(forward_demodulation, [], [f11703, f2822])).
fof(f11703, plain, (~ (h1(op1(e10, e11)) = op2(h1(e10), e24)) | (spl68_493 | ~ spl68_519)), inference(forward_demodulation, [], [f3626, f3759])).
fof(f3626, plain, (~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | spl68_493), inference(avatar_component_clause, [], [f3624])).
fof(f3624, plain, (spl68_493 <=> (h1(op1(e10, e11)) = op2(h1(e10), h1(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl68_493])])).
fof(f11699, plain, (~ spl68_28 | spl68_514), inference(avatar_contradiction_clause, [], [f11698])).
fof(f11698, plain, ($false | (~ spl68_28 | spl68_514)), inference(trivial_inequality_removal, [], [f11697])).
fof(f11697, plain, (~ (h1(e12) = h1(e12)) | (~ spl68_28 | spl68_514)), inference(forward_demodulation, [], [f3710, f1004])).
fof(f1004, plain, ((e12 = op1(e13, e14)) | ~ spl68_28), inference(avatar_component_clause, [], [f1002])).
fof(f1002, plain, (spl68_28 <=> (e12 = op1(e13, e14))), introduced(avatar_definition, [new_symbols(naming, [spl68_28])])).
fof(f3710, plain, (~ (h1(e12) = h1(op1(e13, e14))) | spl68_514), inference(avatar_component_clause, [], [f3708])).
fof(f3708, plain, (spl68_514 <=> (h1(e12) = h1(op1(e13, e14)))), introduced(avatar_definition, [new_symbols(naming, [spl68_514])])).
fof(f11696, plain, (~ spl68_121 | ~ spl68_193 | ~ spl68_372 | spl68_492), inference(avatar_contradiction_clause, [], [f11695])).
fof(f11695, plain, ($false | (~ spl68_121 | ~ spl68_193 | ~ spl68_372 | spl68_492)), inference(subsumption_resolution, [], [f11694, f1747])).
fof(f1747, plain, ((e22 = op2(e22, e22)) | ~ spl68_193), inference(avatar_component_clause, [], [f1745])).
fof(f1745, plain, (spl68_193 <=> (e22 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl68_193])])).
fof(f11694, plain, (~ (e22 = op2(e22, e22)) | (~ spl68_121 | ~ spl68_372 | spl68_492)), inference(forward_demodulation, [], [f11693, f2822])).
fof(f11693, plain, (~ (op2(e22, e22) = h1(e10)) | (~ spl68_121 | ~ spl68_372 | spl68_492)), inference(forward_demodulation, [], [f11692, f1395])).
fof(f1395, plain, ((e10 = op1(e10, e10)) | ~ spl68_121), inference(avatar_component_clause, [], [f1393])).
fof(f1393, plain, (spl68_121 <=> (e10 = op1(e10, e10))), introduced(avatar_definition, [new_symbols(naming, [spl68_121])])).
fof(f11692, plain, (~ (op2(e22, e22) = h1(op1(e10, e10))) | (~ spl68_372 | spl68_492)), inference(forward_demodulation, [], [f3622, f2822])).
fof(f3622, plain, (~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))) | spl68_492), inference(avatar_component_clause, [], [f3620])).
fof(f3620, plain, (spl68_492 <=> (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl68_492])])).
fof(f11620, plain, (~ spl68_54 | ~ spl68_229 | ~ spl68_364 | ~ spl68_375 | spl68_515), inference(avatar_contradiction_clause, [], [f11619])).
fof(f11619, plain, ($false | (~ spl68_54 | ~ spl68_229 | ~ spl68_364 | ~ spl68_375 | spl68_515)), inference(subsumption_resolution, [], [f11618, f1898])).
fof(f11618, plain, (~ (e23 = op2(e21, e20)) | (~ spl68_54 | ~ spl68_364 | ~ spl68_375 | spl68_515)), inference(forward_demodulation, [], [f11617, f2782])).
fof(f11617, plain, (~ (op2(e21, e20) = h1(e13)) | (~ spl68_54 | ~ spl68_375 | spl68_515)), inference(forward_demodulation, [], [f11616, f1113])).
fof(f1113, plain, ((e13 = op1(e12, e14)) | ~ spl68_54), inference(avatar_component_clause, [], [f1111])).
fof(f1111, plain, (spl68_54 <=> (e13 = op1(e12, e14))), introduced(avatar_definition, [new_symbols(naming, [spl68_54])])).
fof(f11616, plain, (~ (op2(e21, e20) = h1(op1(e12, e14))) | (~ spl68_375 | spl68_515)), inference(forward_demodulation, [], [f3714, f2837])).
fof(f3714, plain, (~ (h1(op1(e12, e14)) = op2(h1(e12), e20)) | spl68_515), inference(avatar_component_clause, [], [f3712])).
fof(f3712, plain, (spl68_515 <=> (h1(op1(e12, e14)) = op2(h1(e12), e20))), introduced(avatar_definition, [new_symbols(naming, [spl68_515])])).
fof(f11610, plain, (~ spl68_76 | spl68_516), inference(avatar_contradiction_clause, [], [f11609])).
fof(f11609, plain, ($false | (~ spl68_76 | spl68_516)), inference(trivial_inequality_removal, [], [f11608])).
fof(f11608, plain, (~ (h1(e10) = h1(e10)) | (~ spl68_76 | spl68_516)), inference(forward_demodulation, [], [f3718, f1206])).
fof(f1206, plain, ((e10 = op1(e11, e14)) | ~ spl68_76), inference(avatar_component_clause, [], [f1204])).
fof(f1204, plain, (spl68_76 <=> (e10 = op1(e11, e14))), introduced(avatar_definition, [new_symbols(naming, [spl68_76])])).
fof(f3718, plain, (~ (h1(e10) = h1(op1(e11, e14))) | spl68_516), inference(avatar_component_clause, [], [f3716])).
fof(f3716, plain, (spl68_516 <=> (h1(e10) = h1(op1(e11, e14)))), introduced(avatar_definition, [new_symbols(naming, [spl68_516])])).
fof(f11473, plain, (~ spl68_105 | ~ spl68_201 | ~ spl68_372 | spl68_517), inference(avatar_contradiction_clause, [], [f11472])).
fof(f11472, plain, ($false | (~ spl68_105 | ~ spl68_201 | ~ spl68_372 | spl68_517)), inference(subsumption_resolution, [], [f11471, f1781])).
fof(f1781, plain, ((e20 = op2(e22, e20)) | ~ spl68_201), inference(avatar_component_clause, [], [f1779])).
fof(f1779, plain, (spl68_201 <=> (e20 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl68_201])])).
fof(f11471, plain, (~ (e20 = op2(e22, e20)) | (~ spl68_105 | ~ spl68_372 | spl68_517)), inference(backward_demodulation, [], [f11363, f2822])).
fof(f11363, plain, (~ (e20 = op2(h1(e10), e20)) | (~ spl68_105 | spl68_517)), inference(forward_demodulation, [], [f11362, f738])).
fof(f11362, plain, (~ (h1(e14) = op2(h1(e10), e20)) | (~ spl68_105 | spl68_517)), inference(forward_demodulation, [], [f3722, f1327])).
fof(f1327, plain, ((e14 = op1(e10, e14)) | ~ spl68_105), inference(avatar_component_clause, [], [f1325])).
fof(f1325, plain, (spl68_105 <=> (e14 = op1(e10, e14))), introduced(avatar_definition, [new_symbols(naming, [spl68_105])])).
fof(f3722, plain, (~ (h1(op1(e10, e14)) = op2(h1(e10), e20)) | spl68_517), inference(avatar_component_clause, [], [f3720])).
fof(f3720, plain, (spl68_517 <=> (h1(op1(e10, e14)) = op2(h1(e10), e20))), introduced(avatar_definition, [new_symbols(naming, [spl68_517])])).
fof(f11467, plain, (~ spl68_326 | ~ spl68_32 | ~ spl68_314 | spl68_388), inference(avatar_split_clause, [], [f11466, f2927, f2524, f1019, f2589])).
fof(f11466, plain, (~ (e21 = h4(e11)) | (~ spl68_32 | ~ spl68_314 | spl68_388)), inference(forward_demodulation, [], [f11465, f2525])).
fof(f11465, plain, (~ (h4(e11) = h5(e11)) | (~ spl68_32 | spl68_388)), inference(forward_demodulation, [], [f2929, f1021])).
fof(f11306, plain, (~ spl68_76 | spl68_429), inference(avatar_contradiction_clause, [], [f11305])).
fof(f11305, plain, ($false | (~ spl68_76 | spl68_429)), inference(trivial_inequality_removal, [], [f11304])).
fof(f11304, plain, (~ (h4(e10) = h4(e10)) | (~ spl68_76 | spl68_429)), inference(forward_demodulation, [], [f3184, f1206])).
fof(f3184, plain, (~ (h4(e10) = h4(op1(e11, e14))) | spl68_429), inference(avatar_component_clause, [], [f3182])).
fof(f3182, plain, (spl68_429 <=> (h4(e10) = h4(op1(e11, e14)))), introduced(avatar_definition, [new_symbols(naming, [spl68_429])])).
fof(f11297, plain, (~ spl68_105 | ~ spl68_189 | ~ spl68_322 | spl68_430), inference(avatar_contradiction_clause, [], [f11296])).
fof(f11296, plain, ($false | (~ spl68_105 | ~ spl68_189 | ~ spl68_322 | spl68_430)), inference(subsumption_resolution, [], [f11295, f1730])).
fof(f11295, plain, (~ (e23 = op2(e22, e23)) | (~ spl68_105 | ~ spl68_322 | spl68_430)), inference(forward_demodulation, [], [f11294, f753])).
fof(f753, plain, (e23 = h4(e14)), inference(cnf_transformation, [], [f19])).
fof(f11294, plain, (~ (op2(e22, e23) = h4(e14)) | (~ spl68_105 | ~ spl68_322 | spl68_430)), inference(forward_demodulation, [], [f11293, f1327])).
fof(f11293, plain, (~ (op2(e22, e23) = h4(op1(e10, e14))) | (~ spl68_322 | spl68_430)), inference(forward_demodulation, [], [f3188, f2570])).
fof(f2570, plain, ((e22 = h4(e10)) | ~ spl68_322), inference(avatar_component_clause, [], [f2569])).
fof(f3188, plain, (~ (h4(op1(e10, e14)) = op2(h4(e10), e23)) | spl68_430), inference(avatar_component_clause, [], [f3186])).
fof(f3186, plain, (spl68_430 <=> (h4(op1(e10, e14)) = op2(h4(e10), e23))), introduced(avatar_definition, [new_symbols(naming, [spl68_430])])).
fof(f11292, plain, (~ spl68_54 | ~ spl68_136 | ~ spl68_329 | spl68_428 | ~ spl68_431), inference(avatar_contradiction_clause, [], [f11291])).
fof(f11291, plain, ($false | (~ spl68_54 | ~ spl68_136 | ~ spl68_329 | spl68_428 | ~ spl68_431)), inference(subsumption_resolution, [], [f11290, f1508])).
fof(f11290, plain, (~ (e20 = op2(e24, e23)) | (~ spl68_54 | ~ spl68_329 | spl68_428 | ~ spl68_431)), inference(forward_demodulation, [], [f11289, f2605])).
fof(f2605, plain, ((e20 = h4(e13)) | ~ spl68_329), inference(avatar_component_clause, [], [f2604])).
fof(f2604, plain, (spl68_329 <=> (e20 = h4(e13))), introduced(avatar_definition, [new_symbols(naming, [spl68_329])])).
fof(f11289, plain, (~ (op2(e24, e23) = h4(e13)) | (~ spl68_54 | spl68_428 | ~ spl68_431)), inference(forward_demodulation, [], [f11288, f1113])).
fof(f11288, plain, (~ (op2(e24, e23) = h4(op1(e12, e14))) | (spl68_428 | ~ spl68_431)), inference(forward_demodulation, [], [f3180, f3206])).
fof(f3206, plain, ((e24 = h4(e12)) | ~ spl68_431), inference(avatar_component_clause, [], [f3205])).
fof(f3205, plain, (spl68_431 <=> (e24 = h4(e12))), introduced(avatar_definition, [new_symbols(naming, [spl68_431])])).
fof(f3180, plain, (~ (h4(op1(e12, e14)) = op2(h4(e12), e23)) | spl68_428), inference(avatar_component_clause, [], [f3178])).
fof(f3178, plain, (spl68_428 <=> (h4(op1(e12, e14)) = op2(h4(e12), e23))), introduced(avatar_definition, [new_symbols(naming, [spl68_428])])).
fof(f11287, plain, (~ spl68_28 | spl68_427), inference(avatar_contradiction_clause, [], [f11286])).
fof(f11286, plain, ($false | (~ spl68_28 | spl68_427)), inference(trivial_inequality_removal, [], [f11285])).
fof(f11285, plain, (~ (h4(e12) = h4(e12)) | (~ spl68_28 | spl68_427)), inference(forward_demodulation, [], [f3176, f1004])).
fof(f3176, plain, (~ (h4(e12) = h4(op1(e13, e14))) | spl68_427), inference(avatar_component_clause, [], [f3174])).
fof(f3174, plain, (spl68_427 <=> (h4(e12) = h4(op1(e13, e14)))), introduced(avatar_definition, [new_symbols(naming, [spl68_427])])).
fof(f11284, plain, (~ spl68_19 | spl68_425), inference(avatar_contradiction_clause, [], [f11283])).
fof(f11283, plain, ($false | (~ spl68_19 | spl68_425)), inference(trivial_inequality_removal, [], [f11282])).
fof(f11282, plain, (~ (h4(e13) = h4(e13)) | (~ spl68_19 | spl68_425)), inference(forward_demodulation, [], [f3168, f966])).
fof(f966, plain, ((e13 = op1(e14, e11)) | ~ spl68_19), inference(avatar_component_clause, [], [f964])).
fof(f964, plain, (spl68_19 <=> (e13 = op1(e14, e11))), introduced(avatar_definition, [new_symbols(naming, [spl68_19])])).
fof(f3168, plain, (~ (h4(e13) = h4(op1(e14, e11))) | spl68_425), inference(avatar_component_clause, [], [f3166])).
fof(f3166, plain, (spl68_425 <=> (h4(e13) = h4(op1(e14, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl68_425])])).
fof(f11275, plain, (~ spl68_25 | ~ spl68_169 | ~ spl68_322 | spl68_426), inference(avatar_contradiction_clause, [], [f11274])).
fof(f11274, plain, ($false | (~ spl68_25 | ~ spl68_169 | ~ spl68_322 | spl68_426)), inference(subsumption_resolution, [], [f11273, f1646])).
fof(f11273, plain, (~ (e23 = op2(e23, e22)) | (~ spl68_25 | ~ spl68_322 | spl68_426)), inference(forward_demodulation, [], [f11272, f753])).
fof(f11272, plain, (~ (op2(e23, e22) = h4(e14)) | (~ spl68_25 | ~ spl68_322 | spl68_426)), inference(forward_demodulation, [], [f11271, f991])).
fof(f11271, plain, (~ (op2(e23, e22) = h4(op1(e14, e10))) | (~ spl68_322 | spl68_426)), inference(forward_demodulation, [], [f3172, f2570])).
fof(f3172, plain, (~ (h4(op1(e14, e10)) = op2(e23, h4(e10))) | spl68_426), inference(avatar_component_clause, [], [f3170])).
fof(f3170, plain, (spl68_426 <=> (h4(op1(e14, e10)) = op2(e23, h4(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl68_426])])).
fof(f11270, plain, (~ spl68_11 | ~ spl68_158 | ~ spl68_322 | spl68_424 | ~ spl68_431), inference(avatar_contradiction_clause, [], [f11269])).
fof(f11269, plain, ($false | (~ spl68_11 | ~ spl68_158 | ~ spl68_322 | spl68_424 | ~ spl68_431)), inference(subsumption_resolution, [], [f11268, f1600])).
fof(f11268, plain, (~ (e22 = op2(e23, e24)) | (~ spl68_11 | ~ spl68_322 | spl68_424 | ~ spl68_431)), inference(forward_demodulation, [], [f11267, f2570])).
fof(f11267, plain, (~ (op2(e23, e24) = h4(e10)) | (~ spl68_11 | spl68_424 | ~ spl68_431)), inference(forward_demodulation, [], [f11266, f933])).
fof(f11266, plain, (~ (op2(e23, e24) = h4(op1(e14, e12))) | (spl68_424 | ~ spl68_431)), inference(forward_demodulation, [], [f3164, f3206])).
fof(f3164, plain, (~ (h4(op1(e14, e12)) = op2(e23, h4(e12))) | spl68_424), inference(avatar_component_clause, [], [f3162])).
fof(f3162, plain, (spl68_424 <=> (h4(op1(e14, e12)) = op2(e23, h4(e12)))), introduced(avatar_definition, [new_symbols(naming, [spl68_424])])).
fof(f11265, plain, (~ spl68_2 | spl68_422), inference(avatar_contradiction_clause, [], [f11264])).
fof(f11264, plain, ($false | (~ spl68_2 | spl68_422)), inference(trivial_inequality_removal, [], [f11263])).
fof(f11263, plain, (~ (h4(e11) = h4(e11)) | (~ spl68_2 | spl68_422)), inference(forward_demodulation, [], [f3156, f895])).
fof(f3156, plain, (~ (h4(e11) = h4(op1(e14, e14))) | spl68_422), inference(avatar_component_clause, [], [f3154])).
fof(f3154, plain, (spl68_422 <=> (h4(e11) = h4(op1(e14, e14)))), introduced(avatar_definition, [new_symbols(naming, [spl68_422])])).
fof(f11256, plain, (~ spl68_8 | ~ spl68_180 | ~ spl68_329 | spl68_423 | ~ spl68_431), inference(avatar_contradiction_clause, [], [f11255])).
fof(f11255, plain, ($false | (~ spl68_8 | ~ spl68_180 | ~ spl68_329 | spl68_423 | ~ spl68_431)), inference(subsumption_resolution, [], [f11254, f1692])).
fof(f1692, plain, ((e24 = op2(e23, e20)) | ~ spl68_180), inference(avatar_component_clause, [], [f1690])).
fof(f11254, plain, (~ (e24 = op2(e23, e20)) | (~ spl68_8 | ~ spl68_329 | spl68_423 | ~ spl68_431)), inference(forward_demodulation, [], [f11253, f3206])).
fof(f11253, plain, (~ (op2(e23, e20) = h4(e12)) | (~ spl68_8 | ~ spl68_329 | spl68_423)), inference(forward_demodulation, [], [f11252, f920])).
fof(f11252, plain, (~ (op2(e23, e20) = h4(op1(e14, e13))) | (~ spl68_329 | spl68_423)), inference(forward_demodulation, [], [f3160, f2605])).
fof(f3160, plain, (~ (h4(op1(e14, e13)) = op2(e23, h4(e13))) | spl68_423), inference(avatar_component_clause, [], [f3158])).
fof(f3158, plain, (spl68_423 <=> (h4(op1(e14, e13)) = op2(e23, h4(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl68_423])])).
fof(f11247, plain, (~ spl68_32 | ~ spl68_252 | ~ spl68_326 | ~ spl68_329 | spl68_420), inference(avatar_contradiction_clause, [], [f11246])).
fof(f11246, plain, ($false | (~ spl68_32 | ~ spl68_252 | ~ spl68_326 | ~ spl68_329 | spl68_420)), inference(subsumption_resolution, [], [f11245, f1995])).
fof(f11245, plain, (~ (op2(e20, e20) = e21) | (~ spl68_32 | ~ spl68_326 | ~ spl68_329 | spl68_420)), inference(forward_demodulation, [], [f11244, f2590])).
fof(f11244, plain, (~ (op2(e20, e20) = h4(e11)) | (~ spl68_32 | ~ spl68_329 | spl68_420)), inference(forward_demodulation, [], [f11243, f1021])).
fof(f11243, plain, (~ (op2(e20, e20) = h4(op1(e13, e13))) | (~ spl68_329 | spl68_420)), inference(forward_demodulation, [], [f3148, f2605])).
fof(f3148, plain, (~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | spl68_420), inference(avatar_component_clause, [], [f3146])).
fof(f3146, plain, (spl68_420 <=> (h4(op1(e13, e13)) = op2(h4(e13), h4(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl68_420])])).
fof(f11238, plain, (~ spl68_40 | ~ spl68_234 | ~ spl68_329 | spl68_419 | ~ spl68_431), inference(avatar_contradiction_clause, [], [f11237])).
fof(f11237, plain, ($false | (~ spl68_40 | ~ spl68_234 | ~ spl68_329 | spl68_419 | ~ spl68_431)), inference(subsumption_resolution, [], [f11236, f1919])).
fof(f11236, plain, (~ (e23 = op2(e20, e24)) | (~ spl68_40 | ~ spl68_329 | spl68_419 | ~ spl68_431)), inference(forward_demodulation, [], [f11235, f753])).
fof(f11235, plain, (~ (op2(e20, e24) = h4(e14)) | (~ spl68_40 | ~ spl68_329 | spl68_419 | ~ spl68_431)), inference(forward_demodulation, [], [f11234, f1054])).
fof(f11234, plain, (~ (op2(e20, e24) = h4(op1(e13, e12))) | (~ spl68_329 | spl68_419 | ~ spl68_431)), inference(forward_demodulation, [], [f11233, f2605])).
fof(f11233, plain, (~ (h4(op1(e13, e12)) = op2(h4(e13), e24)) | (spl68_419 | ~ spl68_431)), inference(forward_demodulation, [], [f3144, f3206])).
fof(f3144, plain, (~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | spl68_419), inference(avatar_component_clause, [], [f3142])).
fof(f3142, plain, (spl68_419 <=> (h4(op1(e13, e12)) = op2(h4(e13), h4(e12)))), introduced(avatar_definition, [new_symbols(naming, [spl68_419])])).
fof(f11228, plain, (~ spl68_41 | ~ spl68_248 | ~ spl68_322 | ~ spl68_326 | ~ spl68_329 | spl68_418), inference(avatar_contradiction_clause, [], [f11227])).
fof(f11227, plain, ($false | (~ spl68_41 | ~ spl68_248 | ~ spl68_322 | ~ spl68_326 | ~ spl68_329 | spl68_418)), inference(subsumption_resolution, [], [f11226, f1978])).
fof(f11226, plain, (~ (e22 = op2(e20, e21)) | (~ spl68_41 | ~ spl68_322 | ~ spl68_326 | ~ spl68_329 | spl68_418)), inference(forward_demodulation, [], [f11225, f2570])).
fof(f11225, plain, (~ (op2(e20, e21) = h4(e10)) | (~ spl68_41 | ~ spl68_326 | ~ spl68_329 | spl68_418)), inference(forward_demodulation, [], [f11224, f1059])).
fof(f11224, plain, (~ (op2(e20, e21) = h4(op1(e13, e11))) | (~ spl68_326 | ~ spl68_329 | spl68_418)), inference(forward_demodulation, [], [f11223, f2605])).
fof(f11223, plain, (~ (h4(op1(e13, e11)) = op2(h4(e13), e21)) | (~ spl68_326 | spl68_418)), inference(forward_demodulation, [], [f3140, f2590])).
fof(f3140, plain, (~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | spl68_418), inference(avatar_component_clause, [], [f3138])).
fof(f3138, plain, (spl68_418 <=> (h4(op1(e13, e11)) = op2(h4(e13), h4(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl68_418])])).
fof(f11219, plain, (~ spl68_49 | ~ spl68_241 | ~ spl68_322 | ~ spl68_329 | spl68_417), inference(avatar_contradiction_clause, [], [f11218])).
fof(f11218, plain, ($false | (~ spl68_49 | ~ spl68_241 | ~ spl68_322 | ~ spl68_329 | spl68_417)), inference(subsumption_resolution, [], [f11217, f1949])).
fof(f11217, plain, (~ (e20 = op2(e20, e22)) | (~ spl68_49 | ~ spl68_322 | ~ spl68_329 | spl68_417)), inference(forward_demodulation, [], [f11216, f2605])).
fof(f11216, plain, (~ (op2(e20, e22) = h4(e13)) | (~ spl68_49 | ~ spl68_322 | ~ spl68_329 | spl68_417)), inference(forward_demodulation, [], [f11215, f1092])).
fof(f11215, plain, (~ (op2(e20, e22) = h4(op1(e13, e10))) | (~ spl68_322 | ~ spl68_329 | spl68_417)), inference(forward_demodulation, [], [f11214, f2605])).
fof(f11214, plain, (~ (h4(op1(e13, e10)) = op2(h4(e13), e22)) | (~ spl68_322 | spl68_417)), inference(forward_demodulation, [], [f3136, f2570])).
fof(f3136, plain, (~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | spl68_417), inference(avatar_component_clause, [], [f3134])).
fof(f3134, plain, (spl68_417 <=> (h4(op1(e13, e10)) = op2(h4(e13), h4(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl68_417])])).
fof(f11209, plain, (~ spl68_62 | ~ spl68_132 | ~ spl68_326 | spl68_415 | ~ spl68_431), inference(avatar_contradiction_clause, [], [f11208])).
fof(f11208, plain, ($false | (~ spl68_62 | ~ spl68_132 | ~ spl68_326 | spl68_415 | ~ spl68_431)), inference(subsumption_resolution, [], [f11207, f1491])).
fof(f11207, plain, (~ (e21 = op2(e24, e24)) | (~ spl68_62 | ~ spl68_326 | spl68_415 | ~ spl68_431)), inference(forward_demodulation, [], [f11206, f2590])).
fof(f11206, plain, (~ (op2(e24, e24) = h4(e11)) | (~ spl68_62 | spl68_415 | ~ spl68_431)), inference(forward_demodulation, [], [f11205, f1147])).
fof(f11205, plain, (~ (op2(e24, e24) = h4(op1(e12, e12))) | (spl68_415 | ~ spl68_431)), inference(forward_demodulation, [], [f3128, f3206])).
fof(f3128, plain, (~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | spl68_415), inference(avatar_component_clause, [], [f3126])).
fof(f3126, plain, (spl68_415 <=> (h4(op1(e12, e12)) = op2(h4(e12), h4(e12)))), introduced(avatar_definition, [new_symbols(naming, [spl68_415])])).
fof(f11200, plain, (~ spl68_70 | ~ spl68_149 | ~ spl68_326 | spl68_414 | ~ spl68_431), inference(avatar_contradiction_clause, [], [f11199])).
fof(f11199, plain, ($false | (~ spl68_70 | ~ spl68_149 | ~ spl68_326 | spl68_414 | ~ spl68_431)), inference(subsumption_resolution, [], [f11198, f1562])).
fof(f11198, plain, (~ (e23 = op2(e24, e21)) | (~ spl68_70 | ~ spl68_326 | spl68_414 | ~ spl68_431)), inference(forward_demodulation, [], [f11197, f753])).
fof(f11197, plain, (~ (op2(e24, e21) = h4(e14)) | (~ spl68_70 | ~ spl68_326 | spl68_414 | ~ spl68_431)), inference(forward_demodulation, [], [f11196, f1180])).
fof(f11196, plain, (~ (op2(e24, e21) = h4(op1(e12, e11))) | (~ spl68_326 | spl68_414 | ~ spl68_431)), inference(forward_demodulation, [], [f11195, f3206])).
fof(f11195, plain, (~ (h4(op1(e12, e11)) = op2(h4(e12), e21)) | (~ spl68_326 | spl68_414)), inference(forward_demodulation, [], [f3124, f2590])).
fof(f3124, plain, (~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | spl68_414), inference(avatar_component_clause, [], [f3122])).
fof(f3122, plain, (spl68_414 <=> (h4(op1(e12, e11)) = op2(h4(e12), h4(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl68_414])])).
fof(f11190, plain, (~ spl68_73 | ~ spl68_145 | ~ spl68_322 | spl68_413 | ~ spl68_431), inference(avatar_contradiction_clause, [], [f11189])).
fof(f11189, plain, ($false | (~ spl68_73 | ~ spl68_145 | ~ spl68_322 | spl68_413 | ~ spl68_431)), inference(subsumption_resolution, [], [f11188, f1545])).
fof(f11188, plain, (~ (e24 = op2(e24, e22)) | (~ spl68_73 | ~ spl68_322 | spl68_413 | ~ spl68_431)), inference(forward_demodulation, [], [f11187, f3206])).
fof(f11187, plain, (~ (op2(e24, e22) = h4(e12)) | (~ spl68_73 | ~ spl68_322 | spl68_413 | ~ spl68_431)), inference(forward_demodulation, [], [f11186, f1193])).
fof(f11186, plain, (~ (op2(e24, e22) = h4(op1(e12, e10))) | (~ spl68_322 | spl68_413 | ~ spl68_431)), inference(forward_demodulation, [], [f11185, f3206])).
fof(f11185, plain, (~ (h4(op1(e12, e10)) = op2(h4(e12), e22)) | (~ spl68_322 | spl68_413)), inference(forward_demodulation, [], [f3120, f2570])).
fof(f3120, plain, (~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | spl68_413), inference(avatar_component_clause, [], [f3118])).
fof(f3118, plain, (spl68_413 <=> (h4(op1(e12, e10)) = op2(h4(e12), h4(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl68_413])])).
fof(f11180, plain, (~ spl68_85 | ~ spl68_229 | ~ spl68_326 | ~ spl68_329 | spl68_412), inference(avatar_contradiction_clause, [], [f11179])).
fof(f11179, plain, ($false | (~ spl68_85 | ~ spl68_229 | ~ spl68_326 | ~ spl68_329 | spl68_412)), inference(subsumption_resolution, [], [f11178, f1898])).
fof(f11178, plain, (~ (e23 = op2(e21, e20)) | (~ spl68_85 | ~ spl68_326 | ~ spl68_329 | spl68_412)), inference(forward_demodulation, [], [f11177, f753])).
fof(f11177, plain, (~ (op2(e21, e20) = h4(e14)) | (~ spl68_85 | ~ spl68_326 | ~ spl68_329 | spl68_412)), inference(forward_demodulation, [], [f11176, f1243])).
fof(f11176, plain, (~ (op2(e21, e20) = h4(op1(e11, e13))) | (~ spl68_326 | ~ spl68_329 | spl68_412)), inference(forward_demodulation, [], [f11175, f2590])).
fof(f11175, plain, (~ (h4(op1(e11, e13)) = op2(h4(e11), e20)) | (~ spl68_329 | spl68_412)), inference(forward_demodulation, [], [f3116, f2605])).
fof(f3116, plain, (~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | spl68_412), inference(avatar_component_clause, [], [f3114])).
fof(f3114, plain, (spl68_412 <=> (h4(op1(e11, e13)) = op2(h4(e11), h4(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl68_412])])).
fof(f11171, plain, (~ spl68_89 | ~ spl68_206 | ~ spl68_326 | ~ spl68_329 | spl68_411 | ~ spl68_431), inference(avatar_contradiction_clause, [], [f11170])).
fof(f11170, plain, ($false | (~ spl68_89 | ~ spl68_206 | ~ spl68_326 | ~ spl68_329 | spl68_411 | ~ spl68_431)), inference(subsumption_resolution, [], [f11169, f1802])).
fof(f11169, plain, (~ (e20 = op2(e21, e24)) | (~ spl68_89 | ~ spl68_326 | ~ spl68_329 | spl68_411 | ~ spl68_431)), inference(forward_demodulation, [], [f11168, f2605])).
fof(f11168, plain, (~ (op2(e21, e24) = h4(e13)) | (~ spl68_89 | ~ spl68_326 | spl68_411 | ~ spl68_431)), inference(forward_demodulation, [], [f11167, f1260])).
fof(f11167, plain, (~ (op2(e21, e24) = h4(op1(e11, e12))) | (~ spl68_326 | spl68_411 | ~ spl68_431)), inference(forward_demodulation, [], [f11166, f2590])).
fof(f11166, plain, (~ (h4(op1(e11, e12)) = op2(h4(e11), e24)) | (spl68_411 | ~ spl68_431)), inference(forward_demodulation, [], [f3112, f3206])).
fof(f3112, plain, (~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | spl68_411), inference(avatar_component_clause, [], [f3110])).
fof(f3110, plain, (spl68_411 <=> (h4(op1(e11, e12)) = op2(h4(e11), h4(e12)))), introduced(avatar_definition, [new_symbols(naming, [spl68_411])])).
fof(f11161, plain, (~ spl68_93 | ~ spl68_225 | ~ spl68_326 | spl68_410 | ~ spl68_431), inference(avatar_contradiction_clause, [], [f11160])).
fof(f11160, plain, ($false | (~ spl68_93 | ~ spl68_225 | ~ spl68_326 | spl68_410 | ~ spl68_431)), inference(subsumption_resolution, [], [f11159, f1881])).
fof(f11159, plain, (~ (e24 = op2(e21, e21)) | (~ spl68_93 | ~ spl68_326 | spl68_410 | ~ spl68_431)), inference(forward_demodulation, [], [f11158, f3206])).
fof(f11158, plain, (~ (op2(e21, e21) = h4(e12)) | (~ spl68_93 | ~ spl68_326 | spl68_410)), inference(forward_demodulation, [], [f11157, f1277])).
fof(f11157, plain, (~ (op2(e21, e21) = h4(op1(e11, e11))) | (~ spl68_326 | spl68_410)), inference(forward_demodulation, [], [f3108, f2590])).
fof(f3108, plain, (~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | spl68_410), inference(avatar_component_clause, [], [f3106])).
fof(f3106, plain, (spl68_410 <=> (h4(op1(e11, e11)) = op2(h4(e11), h4(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl68_410])])).
fof(f11152, plain, (~ spl68_97 | ~ spl68_217 | ~ spl68_322 | ~ spl68_326 | spl68_409), inference(avatar_contradiction_clause, [], [f11151])).
fof(f11151, plain, ($false | (~ spl68_97 | ~ spl68_217 | ~ spl68_322 | ~ spl68_326 | spl68_409)), inference(subsumption_resolution, [], [f11150, f1848])).
fof(f11150, plain, (~ (e21 = op2(e21, e22)) | (~ spl68_97 | ~ spl68_322 | ~ spl68_326 | spl68_409)), inference(forward_demodulation, [], [f11149, f2590])).
fof(f11149, plain, (~ (op2(e21, e22) = h4(e11)) | (~ spl68_97 | ~ spl68_322 | ~ spl68_326 | spl68_409)), inference(forward_demodulation, [], [f11148, f1294])).
fof(f11148, plain, (~ (op2(e21, e22) = h4(op1(e11, e10))) | (~ spl68_322 | ~ spl68_326 | spl68_409)), inference(forward_demodulation, [], [f11147, f2590])).
fof(f11147, plain, (~ (h4(op1(e11, e10)) = op2(h4(e11), e22)) | (~ spl68_322 | spl68_409)), inference(forward_demodulation, [], [f3104, f2570])).
fof(f3104, plain, (~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | spl68_409), inference(avatar_component_clause, [], [f3102])).
fof(f3102, plain, (spl68_409 <=> (h4(op1(e11, e10)) = op2(h4(e11), h4(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl68_409])])).
fof(f11142, plain, (~ spl68_109 | ~ spl68_201 | ~ spl68_322 | ~ spl68_329 | spl68_408), inference(avatar_contradiction_clause, [], [f11141])).
fof(f11141, plain, ($false | (~ spl68_109 | ~ spl68_201 | ~ spl68_322 | ~ spl68_329 | spl68_408)), inference(subsumption_resolution, [], [f11140, f1781])).
fof(f11140, plain, (~ (e20 = op2(e22, e20)) | (~ spl68_109 | ~ spl68_322 | ~ spl68_329 | spl68_408)), inference(forward_demodulation, [], [f11139, f2605])).
fof(f11139, plain, (~ (op2(e22, e20) = h4(e13)) | (~ spl68_109 | ~ spl68_322 | ~ spl68_329 | spl68_408)), inference(forward_demodulation, [], [f11138, f1344])).
fof(f11138, plain, (~ (op2(e22, e20) = h4(op1(e10, e13))) | (~ spl68_322 | ~ spl68_329 | spl68_408)), inference(forward_demodulation, [], [f11137, f2570])).
fof(f11137, plain, (~ (h4(op1(e10, e13)) = op2(h4(e10), e20)) | (~ spl68_329 | spl68_408)), inference(forward_demodulation, [], [f3100, f2605])).
fof(f3100, plain, (~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | spl68_408), inference(avatar_component_clause, [], [f3098])).
fof(f3098, plain, (spl68_408 <=> (h4(op1(e10, e13)) = op2(h4(e10), h4(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl68_408])])).
fof(f11136, plain, (~ spl68_117 | ~ spl68_197 | ~ spl68_322 | ~ spl68_326 | spl68_406), inference(avatar_contradiction_clause, [], [f11135])).
fof(f11135, plain, ($false | (~ spl68_117 | ~ spl68_197 | ~ spl68_322 | ~ spl68_326 | spl68_406)), inference(subsumption_resolution, [], [f11134, f1764])).
fof(f11134, plain, (~ (e21 = op2(e22, e21)) | (~ spl68_117 | ~ spl68_322 | ~ spl68_326 | spl68_406)), inference(forward_demodulation, [], [f11133, f2590])).
fof(f11133, plain, (~ (op2(e22, e21) = h4(e11)) | (~ spl68_117 | ~ spl68_322 | ~ spl68_326 | spl68_406)), inference(forward_demodulation, [], [f11132, f1378])).
fof(f11132, plain, (~ (op2(e22, e21) = h4(op1(e10, e11))) | (~ spl68_322 | ~ spl68_326 | spl68_406)), inference(forward_demodulation, [], [f11131, f2570])).
fof(f11131, plain, (~ (h4(op1(e10, e11)) = op2(h4(e10), e21)) | (~ spl68_326 | spl68_406)), inference(forward_demodulation, [], [f3092, f2590])).
fof(f3092, plain, (~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | spl68_406), inference(avatar_component_clause, [], [f3090])).
fof(f3090, plain, (spl68_406 <=> (h4(op1(e10, e11)) = op2(h4(e10), h4(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl68_406])])).
fof(f11001, plain, (~ spl68_56 | ~ spl68_153 | ~ spl68_322 | ~ spl68_329 | spl68_416 | ~ spl68_431), inference(avatar_contradiction_clause, [], [f11000])).
fof(f11000, plain, ($false | (~ spl68_56 | ~ spl68_153 | ~ spl68_322 | ~ spl68_329 | spl68_416 | ~ spl68_431)), inference(subsumption_resolution, [], [f10999, f1579])).
fof(f10999, plain, (~ (e22 = op2(e24, e20)) | (~ spl68_56 | ~ spl68_322 | ~ spl68_329 | spl68_416 | ~ spl68_431)), inference(forward_demodulation, [], [f10998, f2570])).
fof(f10998, plain, (~ (op2(e24, e20) = h4(e10)) | (~ spl68_56 | ~ spl68_329 | spl68_416 | ~ spl68_431)), inference(forward_demodulation, [], [f10997, f1122])).
fof(f10997, plain, (~ (op2(e24, e20) = h4(op1(e12, e13))) | (~ spl68_329 | spl68_416 | ~ spl68_431)), inference(forward_demodulation, [], [f10996, f3206])).
fof(f10996, plain, (~ (h4(op1(e12, e13)) = op2(h4(e12), e20)) | (~ spl68_329 | spl68_416)), inference(forward_demodulation, [], [f3132, f2605])).
fof(f3132, plain, (~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | spl68_416), inference(avatar_component_clause, [], [f3130])).
fof(f3130, plain, (spl68_416 <=> (h4(op1(e12, e13)) = op2(h4(e12), h4(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl68_416])])).
fof(f10795, plain, (~ spl68_141 | ~ spl68_241), inference(avatar_split_clause, [], [f10788, f1947, f1527])).
fof(f1527, plain, (spl68_141 <=> (e20 = op2(e24, e22))), introduced(avatar_definition, [new_symbols(naming, [spl68_141])])).
fof(f10788, plain, (~ (e20 = op2(e24, e22)) | ~ spl68_241), inference(backward_demodulation, [], [f458, f1949])).
fof(f458, plain, ~ (op2(e20, e22) = op2(e24, e22)), inference(cnf_transformation, [], [f8])).
fof(f10780, plain, (spl68_145 | ~ spl68_258), inference(avatar_split_clause, [], [f10768, f2018, f1543])).
fof(f2018, plain, (spl68_258 <=> (e22 = unit2)), introduced(avatar_definition, [new_symbols(naming, [spl68_258])])).
fof(f10768, plain, ((e24 = op2(e24, e22)) | ~ spl68_258), inference(backward_demodulation, [], [f283, f2020])).
fof(f2020, plain, ((e22 = unit2) | ~ spl68_258), inference(avatar_component_clause, [], [f2018])).
fof(f283, plain, (e24 = op2(e24, unit2)), inference(cnf_transformation, [], [f5])).
fof(f5, plain, (((e24 = unit2) | (e23 = unit2) | (e22 = unit2) | (e21 = unit2) | (e20 = unit2)) & (e24 = op2(e24, unit2)) & (e24 = op2(unit2, e24)) & (e23 = op2(e23, unit2)) & (e23 = op2(unit2, e23)) & (e22 = op2(e22, unit2)) & (e22 = op2(unit2, e22)) & (e21 = op2(e21, unit2)) & (e21 = op2(unit2, e21)) & (e20 = op2(e20, unit2)) & (e20 = op2(unit2, e20))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG052+1.p', ax5)).
fof(f10779, plain, (spl68_185 | ~ spl68_258), inference(avatar_split_clause, [], [f10767, f2018, f1711])).
fof(f10767, plain, ((e24 = op2(e22, e24)) | ~ spl68_258), inference(backward_demodulation, [], [f282, f2020])).
fof(f282, plain, (e24 = op2(unit2, e24)), inference(cnf_transformation, [], [f5])).
fof(f10778, plain, (spl68_169 | ~ spl68_258), inference(avatar_split_clause, [], [f10766, f2018, f1644])).
fof(f10766, plain, ((e23 = op2(e23, e22)) | ~ spl68_258), inference(backward_demodulation, [], [f281, f2020])).
fof(f281, plain, (e23 = op2(e23, unit2)), inference(cnf_transformation, [], [f5])).
fof(f10777, plain, (spl68_189 | ~ spl68_258), inference(avatar_split_clause, [], [f10765, f2018, f1728])).
fof(f10765, plain, ((e23 = op2(e22, e23)) | ~ spl68_258), inference(backward_demodulation, [], [f280, f2020])).
fof(f280, plain, (e23 = op2(unit2, e23)), inference(cnf_transformation, [], [f5])).
fof(f10776, plain, (spl68_193 | ~ spl68_258), inference(avatar_split_clause, [], [f10764, f2018, f1745])).
fof(f10764, plain, ((e22 = op2(e22, e22)) | ~ spl68_258), inference(backward_demodulation, [], [f279, f2020])).
fof(f279, plain, (e22 = op2(e22, unit2)), inference(cnf_transformation, [], [f5])).
fof(f10774, plain, (spl68_217 | ~ spl68_258), inference(avatar_split_clause, [], [f10762, f2018, f1846])).
fof(f10762, plain, ((e21 = op2(e21, e22)) | ~ spl68_258), inference(backward_demodulation, [], [f277, f2020])).
fof(f277, plain, (e21 = op2(e21, unit2)), inference(cnf_transformation, [], [f5])).
fof(f10773, plain, (spl68_197 | ~ spl68_258), inference(avatar_split_clause, [], [f10761, f2018, f1762])).
fof(f10761, plain, ((e21 = op2(e22, e21)) | ~ spl68_258), inference(backward_demodulation, [], [f276, f2020])).
fof(f276, plain, (e21 = op2(unit2, e21)), inference(cnf_transformation, [], [f5])).
fof(f10772, plain, (spl68_241 | ~ spl68_258), inference(avatar_split_clause, [], [f10760, f2018, f1947])).
fof(f10760, plain, ((e20 = op2(e20, e22)) | ~ spl68_258), inference(backward_demodulation, [], [f275, f2020])).
fof(f275, plain, (e20 = op2(e20, unit2)), inference(cnf_transformation, [], [f5])).
fof(f10771, plain, (spl68_201 | ~ spl68_258), inference(avatar_split_clause, [], [f10759, f2018, f1779])).
fof(f10759, plain, ((e20 = op2(e22, e20)) | ~ spl68_258), inference(backward_demodulation, [], [f274, f2020])).
fof(f274, plain, (e20 = op2(unit2, e20)), inference(cnf_transformation, [], [f5])).
fof(f10758, plain, (~ spl68_185 | ~ spl68_113 | ~ spl68_322 | spl68_407 | ~ spl68_431), inference(avatar_split_clause, [], [f10753, f3205, f3094, f2569, f1359, f1711])).
fof(f3094, plain, (spl68_407 <=> (h4(op1(e10, e12)) = op2(h4(e10), h4(e12)))), introduced(avatar_definition, [new_symbols(naming, [spl68_407])])).
fof(f10753, plain, (~ (e24 = op2(e22, e24)) | (~ spl68_113 | ~ spl68_322 | spl68_407 | ~ spl68_431)), inference(backward_demodulation, [], [f10643, f2570])).
fof(f10643, plain, (~ (e24 = op2(h4(e10), e24)) | (~ spl68_113 | spl68_407 | ~ spl68_431)), inference(backward_demodulation, [], [f10585, f3206])).
fof(f10585, plain, (~ (h4(e12) = op2(h4(e10), h4(e12))) | (~ spl68_113 | spl68_407)), inference(forward_demodulation, [], [f3096, f1361])).
fof(f3096, plain, (~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | spl68_407), inference(avatar_component_clause, [], [f3094])).
fof(f10755, plain, (~ spl68_193 | ~ spl68_121 | ~ spl68_322 | spl68_405), inference(avatar_split_clause, [], [f10748, f3086, f2569, f1393, f1745])).
fof(f3086, plain, (spl68_405 <=> (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl68_405])])).
fof(f10748, plain, (~ (e22 = op2(e22, e22)) | (~ spl68_121 | ~ spl68_322 | spl68_405)), inference(backward_demodulation, [], [f10584, f2570])).
fof(f10584, plain, (~ (h4(e10) = op2(h4(e10), h4(e10))) | (~ spl68_121 | spl68_405)), inference(forward_demodulation, [], [f3088, f1395])).
fof(f3088, plain, (~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10))) | spl68_405), inference(avatar_component_clause, [], [f3086])).
fof(f10699, plain, (~ spl68_105 | spl68_176 | ~ spl68_367 | ~ spl68_517), inference(avatar_contradiction_clause, [], [f10698])).
fof(f10698, plain, ($false | (~ spl68_105 | spl68_176 | ~ spl68_367 | ~ spl68_517)), inference(subsumption_resolution, [], [f10691, f1675])).
fof(f1675, plain, (~ (e20 = op2(e23, e20)) | spl68_176), inference(avatar_component_clause, [], [f1674])).
fof(f1674, plain, (spl68_176 <=> (e20 = op2(e23, e20))), introduced(avatar_definition, [new_symbols(naming, [spl68_176])])).
fof(f10691, plain, ((e20 = op2(e23, e20)) | (~ spl68_105 | ~ spl68_367 | ~ spl68_517)), inference(backward_demodulation, [], [f10555, f2797])).
fof(f2797, plain, ((e23 = h1(e10)) | ~ spl68_367), inference(avatar_component_clause, [], [f2796])).
fof(f10555, plain, ((e20 = op2(h1(e10), e20)) | (~ spl68_105 | ~ spl68_517)), inference(forward_demodulation, [], [f10554, f738])).
fof(f10554, plain, ((h1(e14) = op2(h1(e10), e20)) | (~ spl68_105 | ~ spl68_517)), inference(forward_demodulation, [], [f3721, f1327])).
fof(f3721, plain, ((h1(op1(e10, e14)) = op2(h1(e10), e20)) | ~ spl68_517), inference(avatar_component_clause, [], [f3720])).
fof(f10635, plain, (spl68_352 | ~ spl68_149 | ~ spl68_490), inference(avatar_split_clause, [], [f10634, f3580, f1560, f2720])).
fof(f10634, plain, ((e23 = h2(e10)) | (~ spl68_149 | ~ spl68_490)), inference(forward_demodulation, [], [f10619, f1562])).
fof(f10619, plain, ((op2(e24, e21) = h2(e10)) | ~ spl68_490), inference(backward_demodulation, [], [f2449, f3581])).
fof(f2449, plain, (h2(e10) = op2(h2(e11), e21)), inference(forward_demodulation, [], [f744, f745])).
fof(f745, plain, (op2(e21, e21) = h2(e11)), inference(cnf_transformation, [], [f17])).
fof(f744, plain, (h2(e10) = op2(op2(e21, e21), e21)), inference(cnf_transformation, [], [f17])).
fof(f10633, plain, (spl68_359 | ~ spl68_206 | ~ spl68_490), inference(avatar_split_clause, [], [f10632, f3580, f1800, f2756])).
fof(f10632, plain, ((e20 = h2(e13)) | (~ spl68_206 | ~ spl68_490)), inference(forward_demodulation, [], [f10618, f1802])).
fof(f10618, plain, ((op2(e21, e24) = h2(e13)) | ~ spl68_490), inference(backward_demodulation, [], [f2448, f3581])).
fof(f2448, plain, (h2(e13) = op2(e21, h2(e11))), inference(backward_demodulation, [], [f747, f745])).
fof(f747, plain, (op2(e21, op2(e21, e21)) = h2(e13)), inference(cnf_transformation, [], [f17])).
fof(f10631, plain, (~ spl68_215 | ~ spl68_490), inference(avatar_split_clause, [], [f10617, f3580, f1837])).
fof(f1837, plain, (spl68_215 <=> (e24 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl68_215])])).
fof(f10617, plain, (~ (e24 = op2(e21, e23)) | ~ spl68_490), inference(backward_demodulation, [], [f2446, f3581])).
fof(f2446, plain, ~ (op2(e21, e23) = h2(e11)), inference(backward_demodulation, [], [f500, f745])).
fof(f500, plain, ~ (op2(e21, e21) = op2(e21, e23)), inference(cnf_transformation, [], [f8])).
fof(f10500, plain, (~ spl68_54 | ~ spl68_184 | ~ spl68_310 | spl68_392), inference(avatar_contradiction_clause, [], [f10499])).
fof(f10499, plain, ($false | (~ spl68_54 | ~ spl68_184 | ~ spl68_310 | spl68_392)), inference(subsumption_resolution, [], [f10498, f1709])).
fof(f1709, plain, ((e23 = op2(e22, e24)) | ~ spl68_184), inference(avatar_component_clause, [], [f1707])).
fof(f1707, plain, (spl68_184 <=> (e23 = op2(e22, e24))), introduced(avatar_definition, [new_symbols(naming, [spl68_184])])).
fof(f10498, plain, (~ (e23 = op2(e22, e24)) | (~ spl68_54 | ~ spl68_310 | spl68_392)), inference(forward_demodulation, [], [f10497, f2472])).
fof(f2472, plain, (e23 = h5(e13)), inference(forward_demodulation, [], [f762, f737])).
fof(f737, plain, (e23 = op2(e24, op2(e24, e24))), inference(cnf_transformation, [], [f15])).
fof(f15, plain, ((e23 = op2(e24, op2(e24, e24))) & (e22 = op2(op2(e24, op2(e24, e24)), e24)) & (e21 = op2(e24, e24)) & (e20 = op2(op2(e24, e24), e24))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG052+1.p', ax15)).
fof(f762, plain, (op2(e24, op2(e24, e24)) = h5(e13)), inference(cnf_transformation, [], [f20])).
fof(f10497, plain, (~ (op2(e22, e24) = h5(e13)) | (~ spl68_54 | ~ spl68_310 | spl68_392)), inference(forward_demodulation, [], [f10496, f1113])).
fof(f10496, plain, (~ (op2(e22, e24) = h5(op1(e12, e14))) | (~ spl68_310 | spl68_392)), inference(forward_demodulation, [], [f2945, f2503])).
fof(f2503, plain, ((e22 = h5(e12)) | ~ spl68_310), inference(avatar_component_clause, [], [f2502])).
fof(f2502, plain, (spl68_310 <=> (e22 = h5(e12))), introduced(avatar_definition, [new_symbols(naming, [spl68_310])])).
fof(f2945, plain, (~ (h5(op1(e12, e14)) = op2(h5(e12), e24)) | spl68_392), inference(avatar_component_clause, [], [f2943])).
fof(f2943, plain, (spl68_392 <=> (h5(op1(e12, e14)) = op2(h5(e12), e24))), introduced(avatar_definition, [new_symbols(naming, [spl68_392])])).
fof(f10489, plain, (~ spl68_117 | ~ spl68_247 | ~ spl68_314 | spl68_401), inference(avatar_contradiction_clause, [], [f10488])).
fof(f10488, plain, ($false | (~ spl68_117 | ~ spl68_247 | ~ spl68_314 | spl68_401)), inference(subsumption_resolution, [], [f10487, f1974])).
fof(f1974, plain, ((e21 = op2(e20, e21)) | ~ spl68_247), inference(avatar_component_clause, [], [f1972])).
fof(f1972, plain, (spl68_247 <=> (e21 = op2(e20, e21))), introduced(avatar_definition, [new_symbols(naming, [spl68_247])])).
fof(f10487, plain, (~ (e21 = op2(e20, e21)) | (~ spl68_117 | ~ spl68_314 | spl68_401)), inference(forward_demodulation, [], [f10486, f2525])).
fof(f10486, plain, (~ (op2(e20, e21) = h5(e11)) | (~ spl68_117 | ~ spl68_314 | spl68_401)), inference(forward_demodulation, [], [f10485, f1378])).
fof(f10485, plain, (~ (op2(e20, e21) = h5(op1(e10, e11))) | (~ spl68_314 | spl68_401)), inference(forward_demodulation, [], [f2981, f2525])).
fof(f2981, plain, (~ (h5(op1(e10, e11)) = op2(e20, h5(e11))) | spl68_401), inference(avatar_component_clause, [], [f2979])).
fof(f2979, plain, (spl68_401 <=> (h5(op1(e10, e11)) = op2(e20, h5(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl68_401])])).
fof(f10484, plain, (~ spl68_11 | ~ spl68_141 | ~ spl68_310 | spl68_384), inference(avatar_contradiction_clause, [], [f10483])).
fof(f10483, plain, ($false | (~ spl68_11 | ~ spl68_141 | ~ spl68_310 | spl68_384)), inference(subsumption_resolution, [], [f10482, f1529])).
fof(f1529, plain, ((e20 = op2(e24, e22)) | ~ spl68_141), inference(avatar_component_clause, [], [f1527])).
fof(f10482, plain, (~ (e20 = op2(e24, e22)) | (~ spl68_11 | ~ spl68_310 | spl68_384)), inference(forward_demodulation, [], [f10481, f2492])).
fof(f2492, plain, (e20 = h5(e10)), inference(forward_demodulation, [], [f2491, f2488])).
fof(f2488, plain, (e20 = op2(h5(e11), e24)), inference(backward_demodulation, [], [f734, f760])).
fof(f760, plain, (op2(e24, e24) = h5(e11)), inference(cnf_transformation, [], [f20])).
fof(f734, plain, (e20 = op2(op2(e24, e24), e24)), inference(cnf_transformation, [], [f15])).
fof(f2491, plain, (h5(e10) = op2(h5(e11), e24)), inference(forward_demodulation, [], [f759, f760])).
fof(f759, plain, (op2(op2(e24, e24), e24) = h5(e10)), inference(cnf_transformation, [], [f20])).
fof(f10481, plain, (~ (op2(e24, e22) = h5(e10)) | (~ spl68_11 | ~ spl68_310 | spl68_384)), inference(forward_demodulation, [], [f10480, f933])).
fof(f10480, plain, (~ (op2(e24, e22) = h5(op1(e14, e12))) | (~ spl68_310 | spl68_384)), inference(forward_demodulation, [], [f2913, f2503])).
fof(f2913, plain, (~ (h5(op1(e14, e12)) = op2(e24, h5(e12))) | spl68_384), inference(avatar_component_clause, [], [f2911])).
fof(f2911, plain, (spl68_384 <=> (h5(op1(e14, e12)) = op2(e24, h5(e12)))), introduced(avatar_definition, [new_symbols(naming, [spl68_384])])).
fof(f10474, plain, (~ spl68_121 | ~ spl68_251 | spl68_402), inference(avatar_contradiction_clause, [], [f10473])).
fof(f10473, plain, ($false | (~ spl68_121 | ~ spl68_251 | spl68_402)), inference(subsumption_resolution, [], [f10472, f8897])).
fof(f8897, plain, ((e20 = h1(e11)) | ~ spl68_251), inference(backward_demodulation, [], [f740, f1991])).
fof(f1991, plain, ((e20 = op2(e20, e20)) | ~ spl68_251), inference(avatar_component_clause, [], [f1989])).
fof(f1989, plain, (spl68_251 <=> (e20 = op2(e20, e20))), introduced(avatar_definition, [new_symbols(naming, [spl68_251])])).
fof(f10472, plain, (~ (e20 = h1(e11)) | (~ spl68_121 | spl68_402)), inference(forward_demodulation, [], [f10471, f2492])).
fof(f10471, plain, (~ (h1(e11) = h5(e10)) | (~ spl68_121 | spl68_402)), inference(forward_demodulation, [], [f2985, f1395])).
fof(f10470, plain, (~ spl68_40 | ~ spl68_85 | ~ spl68_170 | ~ spl68_215 | ~ spl68_310 | ~ spl68_314 | spl68_389 | ~ spl68_396), inference(avatar_contradiction_clause, [], [f10469])).
fof(f10469, plain, ($false | (~ spl68_40 | ~ spl68_85 | ~ spl68_170 | ~ spl68_215 | ~ spl68_310 | ~ spl68_314 | spl68_389 | ~ spl68_396)), inference(subsumption_resolution, [], [f10468, f1650])).
fof(f1650, plain, ((e24 = op2(e23, e22)) | ~ spl68_170), inference(avatar_component_clause, [], [f1648])).
fof(f1648, plain, (spl68_170 <=> (e24 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl68_170])])).
fof(f10468, plain, (~ (e24 = op2(e23, e22)) | (~ spl68_40 | ~ spl68_85 | ~ spl68_215 | ~ spl68_310 | ~ spl68_314 | spl68_389 | ~ spl68_396)), inference(forward_demodulation, [], [f10467, f10216])).
fof(f10216, plain, ((e24 = h5(e14)) | (~ spl68_85 | ~ spl68_215 | ~ spl68_314 | ~ spl68_396)), inference(forward_demodulation, [], [f10159, f1839])).
fof(f1839, plain, ((e24 = op2(e21, e23)) | ~ spl68_215), inference(avatar_component_clause, [], [f1837])).
fof(f10159, plain, ((op2(e21, e23) = h5(e14)) | (~ spl68_85 | ~ spl68_314 | ~ spl68_396)), inference(forward_demodulation, [], [f10158, f1243])).
fof(f10158, plain, ((op2(e21, e23) = h5(op1(e11, e13))) | (~ spl68_314 | ~ spl68_396)), inference(forward_demodulation, [], [f2960, f2525])).
fof(f2960, plain, ((h5(op1(e11, e13)) = op2(h5(e11), e23)) | ~ spl68_396), inference(avatar_component_clause, [], [f2959])).
fof(f2959, plain, (spl68_396 <=> (h5(op1(e11, e13)) = op2(h5(e11), e23))), introduced(avatar_definition, [new_symbols(naming, [spl68_396])])).
fof(f10467, plain, (~ (op2(e23, e22) = h5(e14)) | (~ spl68_40 | ~ spl68_310 | spl68_389)), inference(forward_demodulation, [], [f10466, f1054])).
fof(f10466, plain, (~ (op2(e23, e22) = h5(op1(e13, e12))) | (~ spl68_310 | spl68_389)), inference(forward_demodulation, [], [f2933, f2503])).
fof(f2933, plain, (~ (h5(op1(e13, e12)) = op2(e23, h5(e12))) | spl68_389), inference(avatar_component_clause, [], [f2931])).
fof(f2931, plain, (spl68_389 <=> (h5(op1(e13, e12)) = op2(e23, h5(e12)))), introduced(avatar_definition, [new_symbols(naming, [spl68_389])])).
fof(f10460, plain, (~ spl68_73 | ~ spl68_203 | ~ spl68_310 | spl68_394), inference(avatar_contradiction_clause, [], [f10459])).
fof(f10459, plain, ($false | (~ spl68_73 | ~ spl68_203 | ~ spl68_310 | spl68_394)), inference(subsumption_resolution, [], [f10458, f1789])).
fof(f1789, plain, ((e22 = op2(e22, e20)) | ~ spl68_203), inference(avatar_component_clause, [], [f1787])).
fof(f1787, plain, (spl68_203 <=> (e22 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl68_203])])).
fof(f10458, plain, (~ (e22 = op2(e22, e20)) | (~ spl68_73 | ~ spl68_310 | spl68_394)), inference(forward_demodulation, [], [f10457, f2503])).
fof(f10457, plain, (~ (op2(e22, e20) = h5(e12)) | (~ spl68_73 | ~ spl68_310 | spl68_394)), inference(forward_demodulation, [], [f10456, f1193])).
fof(f10456, plain, (~ (op2(e22, e20) = h5(op1(e12, e10))) | (~ spl68_310 | spl68_394)), inference(forward_demodulation, [], [f2953, f2503])).
fof(f2953, plain, (~ (h5(op1(e12, e10)) = op2(h5(e12), e20)) | spl68_394), inference(avatar_component_clause, [], [f2951])).
fof(f2951, plain, (spl68_394 <=> (h5(op1(e12, e10)) = op2(h5(e12), e20))), introduced(avatar_definition, [new_symbols(naming, [spl68_394])])).
fof(f10455, plain, (~ spl68_25 | ~ spl68_85 | ~ spl68_155 | ~ spl68_215 | ~ spl68_314 | spl68_386 | ~ spl68_396), inference(avatar_contradiction_clause, [], [f10454])).
fof(f10454, plain, ($false | (~ spl68_25 | ~ spl68_85 | ~ spl68_155 | ~ spl68_215 | ~ spl68_314 | spl68_386 | ~ spl68_396)), inference(subsumption_resolution, [], [f10453, f1587])).
fof(f1587, plain, ((e24 = op2(e24, e20)) | ~ spl68_155), inference(avatar_component_clause, [], [f1585])).
fof(f1585, plain, (spl68_155 <=> (e24 = op2(e24, e20))), introduced(avatar_definition, [new_symbols(naming, [spl68_155])])).
fof(f10453, plain, (~ (e24 = op2(e24, e20)) | (~ spl68_25 | ~ spl68_85 | ~ spl68_215 | ~ spl68_314 | spl68_386 | ~ spl68_396)), inference(forward_demodulation, [], [f10452, f10216])).
fof(f10452, plain, (~ (op2(e24, e20) = h5(e14)) | (~ spl68_25 | spl68_386)), inference(forward_demodulation, [], [f2921, f991])).
fof(f2921, plain, (~ (op2(e24, e20) = h5(op1(e14, e10))) | spl68_386), inference(avatar_component_clause, [], [f2919])).
fof(f2919, plain, (spl68_386 <=> (op2(e24, e20) = h5(op1(e14, e10)))), introduced(avatar_definition, [new_symbols(naming, [spl68_386])])).
fof(f10445, plain, (~ spl68_97 | ~ spl68_227 | ~ spl68_314 | spl68_397), inference(avatar_contradiction_clause, [], [f10444])).
fof(f10444, plain, ($false | (~ spl68_97 | ~ spl68_227 | ~ spl68_314 | spl68_397)), inference(subsumption_resolution, [], [f10443, f1890])).
fof(f1890, plain, ((e21 = op2(e21, e20)) | ~ spl68_227), inference(avatar_component_clause, [], [f1888])).
fof(f1888, plain, (spl68_227 <=> (e21 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl68_227])])).
fof(f10443, plain, (~ (e21 = op2(e21, e20)) | (~ spl68_97 | ~ spl68_314 | spl68_397)), inference(forward_demodulation, [], [f10442, f2525])).
fof(f10442, plain, (~ (op2(e21, e20) = h5(e11)) | (~ spl68_97 | ~ spl68_314 | spl68_397)), inference(forward_demodulation, [], [f10441, f1294])).
fof(f10441, plain, (~ (op2(e21, e20) = h5(op1(e11, e10))) | (~ spl68_314 | spl68_397)), inference(forward_demodulation, [], [f2965, f2525])).
fof(f2965, plain, (~ (h5(op1(e11, e10)) = op2(h5(e11), e20)) | spl68_397), inference(avatar_component_clause, [], [f2963])).
fof(f2963, plain, (spl68_397 <=> (h5(op1(e11, e10)) = op2(h5(e11), e20))), introduced(avatar_definition, [new_symbols(naming, [spl68_397])])).
fof(f10440, plain, (~ spl68_62 | ~ spl68_192 | ~ spl68_310 | ~ spl68_314 | spl68_381), inference(avatar_contradiction_clause, [], [f10439])).
fof(f10439, plain, ($false | (~ spl68_62 | ~ spl68_192 | ~ spl68_310 | ~ spl68_314 | spl68_381)), inference(subsumption_resolution, [], [f10438, f1743])).
fof(f1743, plain, ((e21 = op2(e22, e22)) | ~ spl68_192), inference(avatar_component_clause, [], [f1741])).
fof(f1741, plain, (spl68_192 <=> (e21 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl68_192])])).
fof(f10438, plain, (~ (e21 = op2(e22, e22)) | (~ spl68_62 | ~ spl68_310 | ~ spl68_314 | spl68_381)), inference(forward_demodulation, [], [f10437, f2525])).
fof(f10437, plain, (~ (op2(e22, e22) = h5(e11)) | (~ spl68_62 | ~ spl68_310 | spl68_381)), inference(forward_demodulation, [], [f10436, f1147])).
fof(f10436, plain, (~ (op2(e22, e22) = h5(op1(e12, e12))) | (~ spl68_310 | spl68_381)), inference(forward_demodulation, [], [f2901, f2503])).
fof(f2901, plain, (~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | spl68_381), inference(avatar_component_clause, [], [f2899])).
fof(f2899, plain, (spl68_381 <=> (h5(op1(e12, e12)) = op2(h5(e12), h5(e12)))), introduced(avatar_definition, [new_symbols(naming, [spl68_381])])).
fof(f10429, plain, (~ spl68_85 | ~ spl68_105 | ~ spl68_215 | ~ spl68_235 | ~ spl68_314 | ~ spl68_396 | spl68_398), inference(avatar_contradiction_clause, [], [f10428])).
fof(f10428, plain, ($false | (~ spl68_85 | ~ spl68_105 | ~ spl68_215 | ~ spl68_235 | ~ spl68_314 | ~ spl68_396 | spl68_398)), inference(subsumption_resolution, [], [f10427, f1923])).
fof(f1923, plain, ((e24 = op2(e20, e24)) | ~ spl68_235), inference(avatar_component_clause, [], [f1921])).
fof(f1921, plain, (spl68_235 <=> (e24 = op2(e20, e24))), introduced(avatar_definition, [new_symbols(naming, [spl68_235])])).
fof(f10427, plain, (~ (e24 = op2(e20, e24)) | (~ spl68_85 | ~ spl68_105 | ~ spl68_215 | ~ spl68_314 | ~ spl68_396 | spl68_398)), inference(forward_demodulation, [], [f10426, f10216])).
fof(f10426, plain, (~ (op2(e20, e24) = h5(e14)) | (~ spl68_105 | spl68_398)), inference(forward_demodulation, [], [f2969, f1327])).
fof(f2969, plain, (~ (op2(e20, e24) = h5(op1(e10, e14))) | spl68_398), inference(avatar_component_clause, [], [f2967])).
fof(f2967, plain, (spl68_398 <=> (op2(e20, e24) = h5(op1(e10, e14)))), introduced(avatar_definition, [new_symbols(naming, [spl68_398])])).
fof(f10425, plain, (~ spl68_70 | ~ spl68_85 | ~ spl68_200 | ~ spl68_215 | ~ spl68_310 | ~ spl68_314 | spl68_380 | ~ spl68_396), inference(avatar_contradiction_clause, [], [f10424])).
fof(f10424, plain, ($false | (~ spl68_70 | ~ spl68_85 | ~ spl68_200 | ~ spl68_215 | ~ spl68_310 | ~ spl68_314 | spl68_380 | ~ spl68_396)), inference(subsumption_resolution, [], [f10423, f1776])).
fof(f1776, plain, ((e24 = op2(e22, e21)) | ~ spl68_200), inference(avatar_component_clause, [], [f1774])).
fof(f1774, plain, (spl68_200 <=> (e24 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl68_200])])).
fof(f10423, plain, (~ (e24 = op2(e22, e21)) | (~ spl68_70 | ~ spl68_85 | ~ spl68_215 | ~ spl68_310 | ~ spl68_314 | spl68_380 | ~ spl68_396)), inference(forward_demodulation, [], [f10422, f10216])).
fof(f10422, plain, (~ (op2(e22, e21) = h5(e14)) | (~ spl68_70 | ~ spl68_310 | ~ spl68_314 | spl68_380)), inference(forward_demodulation, [], [f10421, f1180])).
fof(f10421, plain, (~ (op2(e22, e21) = h5(op1(e12, e11))) | (~ spl68_310 | ~ spl68_314 | spl68_380)), inference(forward_demodulation, [], [f10420, f2503])).
fof(f10420, plain, (~ (h5(op1(e12, e11)) = op2(h5(e12), e21)) | (~ spl68_314 | spl68_380)), inference(forward_demodulation, [], [f2897, f2525])).
fof(f2897, plain, (~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | spl68_380), inference(avatar_component_clause, [], [f2895])).
fof(f2895, plain, (spl68_380 <=> (h5(op1(e12, e11)) = op2(h5(e12), h5(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl68_380])])).
fof(f10412, plain, (~ spl68_113 | ~ spl68_243 | ~ spl68_310 | spl68_400), inference(avatar_contradiction_clause, [], [f10411])).
fof(f10411, plain, ($false | (~ spl68_113 | ~ spl68_243 | ~ spl68_310 | spl68_400)), inference(subsumption_resolution, [], [f10410, f1957])).
fof(f1957, plain, ((e22 = op2(e20, e22)) | ~ spl68_243), inference(avatar_component_clause, [], [f1955])).
fof(f1955, plain, (spl68_243 <=> (e22 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl68_243])])).
fof(f10410, plain, (~ (e22 = op2(e20, e22)) | (~ spl68_113 | ~ spl68_310 | spl68_400)), inference(forward_demodulation, [], [f10409, f2503])).
fof(f10409, plain, (~ (op2(e20, e22) = h5(e12)) | (~ spl68_113 | ~ spl68_310 | spl68_400)), inference(forward_demodulation, [], [f10408, f1361])).
fof(f10408, plain, (~ (op2(e20, e22) = h5(op1(e10, e12))) | (~ spl68_310 | spl68_400)), inference(forward_demodulation, [], [f2977, f2503])).
fof(f2977, plain, (~ (h5(op1(e10, e12)) = op2(e20, h5(e12))) | spl68_400), inference(avatar_component_clause, [], [f2975])).
fof(f2975, plain, (spl68_400 <=> (h5(op1(e10, e12)) = op2(e20, h5(e12)))), introduced(avatar_definition, [new_symbols(naming, [spl68_400])])).
fof(f10407, plain, (~ spl68_89 | ~ spl68_219 | ~ spl68_310 | ~ spl68_314 | spl68_379), inference(avatar_contradiction_clause, [], [f10406])).
fof(f10406, plain, ($false | (~ spl68_89 | ~ spl68_219 | ~ spl68_310 | ~ spl68_314 | spl68_379)), inference(subsumption_resolution, [], [f10405, f1856])).
fof(f1856, plain, ((e23 = op2(e21, e22)) | ~ spl68_219), inference(avatar_component_clause, [], [f1854])).
fof(f1854, plain, (spl68_219 <=> (e23 = op2(e21, e22))), introduced(avatar_definition, [new_symbols(naming, [spl68_219])])).
fof(f10405, plain, (~ (e23 = op2(e21, e22)) | (~ spl68_89 | ~ spl68_310 | ~ spl68_314 | spl68_379)), inference(forward_demodulation, [], [f10404, f2472])).
fof(f10404, plain, (~ (op2(e21, e22) = h5(e13)) | (~ spl68_89 | ~ spl68_310 | ~ spl68_314 | spl68_379)), inference(forward_demodulation, [], [f10403, f1260])).
fof(f10403, plain, (~ (op2(e21, e22) = h5(op1(e11, e12))) | (~ spl68_310 | ~ spl68_314 | spl68_379)), inference(forward_demodulation, [], [f10402, f2525])).
fof(f10402, plain, (~ (h5(op1(e11, e12)) = op2(h5(e11), e22)) | (~ spl68_310 | spl68_379)), inference(forward_demodulation, [], [f2893, f2503])).
fof(f2893, plain, (~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | spl68_379), inference(avatar_component_clause, [], [f2891])).
fof(f2891, plain, (spl68_379 <=> (h5(op1(e11, e12)) = op2(h5(e11), h5(e12)))), introduced(avatar_definition, [new_symbols(naming, [spl68_379])])).
fof(f10389, plain, (~ spl68_31 | ~ spl68_41), inference(avatar_split_clause, [], [f10381, f1057, f1015])).
fof(f1015, plain, (spl68_31 <=> (e10 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl68_31])])).
fof(f10381, plain, (~ (e10 = op1(e13, e13)) | ~ spl68_41), inference(backward_demodulation, [], [f420, f1059])).
fof(f420, plain, ~ (op1(e13, e11) = op1(e13, e13)), inference(cnf_transformation, [], [f7])).
fof(f10378, plain, (~ spl68_37 | ~ spl68_62), inference(avatar_split_clause, [], [f10372, f1145, f1040])).
fof(f1040, plain, (spl68_37 <=> (e11 = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl68_37])])).
fof(f10372, plain, (~ (e11 = op1(e13, e12)) | ~ spl68_62), inference(backward_demodulation, [], [f362, f1147])).
fof(f362, plain, ~ (op1(e12, e12) = op1(e13, e12)), inference(cnf_transformation, [], [f7])).
fof(f10317, plain, (~ spl68_56 | ~ spl68_186 | ~ spl68_310 | spl68_393), inference(avatar_contradiction_clause, [], [f10316])).
fof(f10316, plain, ($false | (~ spl68_56 | ~ spl68_186 | ~ spl68_310 | spl68_393)), inference(subsumption_resolution, [], [f10315, f1718])).
fof(f1718, plain, ((e20 = op2(e22, e23)) | ~ spl68_186), inference(avatar_component_clause, [], [f1716])).
fof(f1716, plain, (spl68_186 <=> (e20 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl68_186])])).
fof(f10315, plain, (~ (e20 = op2(e22, e23)) | (~ spl68_56 | ~ spl68_310 | spl68_393)), inference(forward_demodulation, [], [f10314, f2492])).
fof(f10314, plain, (~ (op2(e22, e23) = h5(e10)) | (~ spl68_56 | ~ spl68_310 | spl68_393)), inference(forward_demodulation, [], [f10283, f1122])).
fof(f10283, plain, (~ (op2(e22, e23) = h5(op1(e12, e13))) | (~ spl68_310 | spl68_393)), inference(forward_demodulation, [], [f2949, f2503])).
fof(f2949, plain, (~ (h5(op1(e12, e13)) = op2(h5(e12), e23)) | spl68_393), inference(avatar_component_clause, [], [f2947])).
fof(f2947, plain, (spl68_393 <=> (h5(op1(e12, e13)) = op2(h5(e12), e23))), introduced(avatar_definition, [new_symbols(naming, [spl68_393])])).
fof(f10296, plain, (~ spl68_116 | ~ spl68_117), inference(avatar_contradiction_clause, [], [f10295])).
fof(f10295, plain, ($false | (~ spl68_116 | ~ spl68_117)), inference(subsumption_resolution, [], [f10294, f535])).
fof(f535, plain, ~ (e10 = e11), inference(cnf_transformation, [], [f9])).
fof(f9, plain, (~ (e13 = e14) & ~ (e12 = e14) & ~ (e12 = e13) & ~ (e11 = e14) & ~ (e11 = e13) & ~ (e11 = e12) & ~ (e10 = e14) & ~ (e10 = e13) & ~ (e10 = e12) & ~ (e10 = e11)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG052+1.p', ax9)).
fof(f10294, plain, ((e10 = e11) | (~ spl68_116 | ~ spl68_117)), inference(backward_demodulation, [], [f1378, f1374])).
fof(f1374, plain, ((e10 = op1(e10, e11)) | ~ spl68_116), inference(avatar_component_clause, [], [f1372])).
fof(f1372, plain, (spl68_116 <=> (e10 = op1(e10, e11))), introduced(avatar_definition, [new_symbols(naming, [spl68_116])])).
fof(f10293, plain, (~ spl68_121 | ~ spl68_124), inference(avatar_contradiction_clause, [], [f10292])).
fof(f10292, plain, ($false | (~ spl68_121 | ~ spl68_124)), inference(subsumption_resolution, [], [f10291, f537])).
fof(f537, plain, ~ (e10 = e13), inference(cnf_transformation, [], [f9])).
fof(f10291, plain, ((e10 = e13) | (~ spl68_121 | ~ spl68_124)), inference(forward_demodulation, [], [f1407, f1395])).
fof(f1407, plain, ((op1(e10, e10) = e13) | ~ spl68_124), inference(avatar_component_clause, [], [f1405])).
fof(f1405, plain, (spl68_124 <=> (op1(e10, e10) = e13)), introduced(avatar_definition, [new_symbols(naming, [spl68_124])])).
fof(f10041, plain, (~ spl68_36 | ~ spl68_11), inference(avatar_split_clause, [], [f10040, f931, f1036])).
fof(f1036, plain, (spl68_36 <=> (e10 = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl68_36])])).
fof(f10040, plain, (~ (e10 = op1(e13, e12)) | ~ spl68_11), inference(forward_demodulation, [], [f364, f933])).
fof(f364, plain, ~ (op1(e13, e12) = op1(e14, e12)), inference(cnf_transformation, [], [f7])).
fof(f10036, plain, (~ spl68_8 | ~ spl68_138 | ~ spl68_310 | spl68_383), inference(avatar_contradiction_clause, [], [f10035])).
fof(f10035, plain, ($false | (~ spl68_8 | ~ spl68_138 | ~ spl68_310 | spl68_383)), inference(subsumption_resolution, [], [f10028, f2503])).
fof(f10028, plain, (~ (e22 = h5(e12)) | (~ spl68_8 | ~ spl68_138 | spl68_383)), inference(backward_demodulation, [], [f9420, f920])).
fof(f9420, plain, (~ (e22 = h5(op1(e14, e13))) | (~ spl68_138 | spl68_383)), inference(forward_demodulation, [], [f2909, f1516])).
fof(f1516, plain, ((e22 = op2(e24, e23)) | ~ spl68_138), inference(avatar_component_clause, [], [f1514])).
fof(f1514, plain, (spl68_138 <=> (e22 = op2(e24, e23))), introduced(avatar_definition, [new_symbols(naming, [spl68_138])])).
fof(f2909, plain, (~ (op2(e24, e23) = h5(op1(e14, e13))) | spl68_383), inference(avatar_component_clause, [], [f2907])).
fof(f2907, plain, (spl68_383 <=> (op2(e24, e23) = h5(op1(e14, e13)))), introduced(avatar_definition, [new_symbols(naming, [spl68_383])])).
fof(f10018, plain, (~ spl68_6 | ~ spl68_11), inference(avatar_split_clause, [], [f10010, f931, f910])).
fof(f910, plain, (spl68_6 <=> (e10 = op1(e14, e13))), introduced(avatar_definition, [new_symbols(naming, [spl68_6])])).
fof(f10010, plain, (~ (e10 = op1(e14, e13)) | ~ spl68_11), inference(backward_demodulation, [], [f432, f933])).
fof(f432, plain, ~ (op1(e14, e12) = op1(e14, e13)), inference(cnf_transformation, [], [f7])).
fof(f10007, plain, (~ spl68_15 | ~ spl68_25), inference(avatar_split_clause, [], [f10003, f989, f947])).
fof(f10003, plain, (~ (e14 = op1(e14, e12)) | ~ spl68_25), inference(backward_demodulation, [], [f426, f991])).
fof(f426, plain, ~ (op1(e14, e10) = op1(e14, e12)), inference(cnf_transformation, [], [f7])).
fof(f9894, plain, (~ spl68_63 | ~ spl68_73), inference(avatar_split_clause, [], [f9888, f1191, f1149])).
fof(f9888, plain, (~ (e12 = op1(e12, e12)) | ~ spl68_73), inference(backward_demodulation, [], [f406, f1193])).
fof(f406, plain, ~ (op1(e12, e10) = op1(e12, e12)), inference(cnf_transformation, [], [f7])).
fof(f9885, plain, (~ spl68_39 | ~ spl68_89), inference(avatar_split_clause, [], [f9880, f1258, f1048])).
fof(f9880, plain, (~ (e13 = op1(e13, e12)) | ~ spl68_89), inference(backward_demodulation, [], [f360, f1260])).
fof(f360, plain, ~ (op1(e11, e12) = op1(e13, e12)), inference(cnf_transformation, [], [f7])).
fof(f9874, plain, (~ spl68_87 | ~ spl68_97), inference(avatar_split_clause, [], [f9868, f1292, f1250])).
fof(f9868, plain, (~ (e11 = op1(e11, e12)) | ~ spl68_97), inference(backward_demodulation, [], [f396, f1294])).
fof(f396, plain, ~ (op1(e11, e10) = op1(e11, e12)), inference(cnf_transformation, [], [f7])).
fof(f9863, plain, (~ spl68_55 | ~ spl68_105), inference(avatar_split_clause, [], [f9859, f1325, f1115])).
fof(f9859, plain, (~ (e14 = op1(e12, e14)) | ~ spl68_105), inference(backward_demodulation, [], [f376, f1327])).
fof(f376, plain, ~ (op1(e10, e14) = op1(e12, e14)), inference(cnf_transformation, [], [f7])).
fof(f9830, plain, (~ spl68_67 | ~ spl68_117), inference(avatar_split_clause, [], [f9819, f1376, f1166])).
fof(f9819, plain, (~ (e11 = op1(e12, e11)) | ~ spl68_117), inference(backward_demodulation, [], [f346, f1378])).
fof(f346, plain, ~ (op1(e10, e11) = op1(e12, e11)), inference(cnf_transformation, [], [f7])).
fof(f9817, plain, (~ spl68_111 | ~ spl68_121), inference(avatar_split_clause, [], [f9806, f1393, f1351])).
fof(f9806, plain, (~ (e10 = op1(e10, e12)) | ~ spl68_121), inference(backward_demodulation, [], [f386, f1395])).
fof(f386, plain, ~ (op1(e10, e10) = op1(e10, e12)), inference(cnf_transformation, [], [f7])).
fof(f9801, plain, (spl68_25 | ~ spl68_126), inference(avatar_split_clause, [], [f9785, f1414, f989])).
fof(f1414, plain, (spl68_126 <=> (e10 = unit1)), introduced(avatar_definition, [new_symbols(naming, [spl68_126])])).
fof(f9785, plain, ((e14 = op1(e14, e10)) | ~ spl68_126), inference(backward_demodulation, [], [f197, f1416])).
fof(f1416, plain, ((e10 = unit1) | ~ spl68_126), inference(avatar_component_clause, [], [f1414])).
fof(f9800, plain, (spl68_105 | ~ spl68_126), inference(avatar_split_clause, [], [f9784, f1414, f1325])).
fof(f9784, plain, ((e14 = op1(e10, e14)) | ~ spl68_126), inference(backward_demodulation, [], [f196, f1416])).
fof(f9799, plain, (~ spl68_47 | ~ spl68_126), inference(avatar_contradiction_clause, [], [f9798])).
fof(f9798, plain, ($false | (~ spl68_47 | ~ spl68_126)), inference(subsumption_resolution, [], [f9797, f540])).
fof(f540, plain, ~ (e11 = e13), inference(cnf_transformation, [], [f9])).
fof(f9797, plain, ((e11 = e13) | (~ spl68_47 | ~ spl68_126)), inference(forward_demodulation, [], [f9783, f1084])).
fof(f9783, plain, ((e13 = op1(e13, e10)) | ~ spl68_126), inference(backward_demodulation, [], [f195, f1416])).
fof(f9796, plain, (spl68_109 | ~ spl68_126), inference(avatar_split_clause, [], [f9782, f1414, f1342])).
fof(f9782, plain, ((e13 = op1(e10, e13)) | ~ spl68_126), inference(backward_demodulation, [], [f194, f1416])).
fof(f9795, plain, (spl68_73 | ~ spl68_126), inference(avatar_split_clause, [], [f9781, f1414, f1191])).
fof(f9781, plain, ((e12 = op1(e12, e10)) | ~ spl68_126), inference(backward_demodulation, [], [f193, f1416])).
fof(f9794, plain, (spl68_113 | ~ spl68_126), inference(avatar_split_clause, [], [f9780, f1414, f1359])).
fof(f9780, plain, ((e12 = op1(e10, e12)) | ~ spl68_126), inference(backward_demodulation, [], [f192, f1416])).
fof(f192, plain, (e12 = op1(unit1, e12)), inference(cnf_transformation, [], [f2])).
fof(f9793, plain, (spl68_97 | ~ spl68_126), inference(avatar_split_clause, [], [f9779, f1414, f1292])).
fof(f9779, plain, ((e11 = op1(e11, e10)) | ~ spl68_126), inference(backward_demodulation, [], [f191, f1416])).
fof(f9792, plain, (spl68_117 | ~ spl68_126), inference(avatar_split_clause, [], [f9778, f1414, f1376])).
fof(f9778, plain, ((e11 = op1(e10, e11)) | ~ spl68_126), inference(backward_demodulation, [], [f190, f1416])).
fof(f9791, plain, (spl68_121 | ~ spl68_126), inference(avatar_split_clause, [], [f9777, f1414, f1393])).
fof(f9777, plain, ((e10 = op1(e10, e10)) | ~ spl68_126), inference(backward_demodulation, [], [f189, f1416])).
fof(f9775, plain, (~ spl68_339 | ~ spl68_450), inference(avatar_contradiction_clause, [], [f9774])).
fof(f9774, plain, ($false | (~ spl68_339 | ~ spl68_450)), inference(subsumption_resolution, [], [f9773, f551])).
fof(f551, plain, ~ (e21 = e24), inference(cnf_transformation, [], [f10])).
fof(f10, plain, (~ (e23 = e24) & ~ (e22 = e24) & ~ (e22 = e23) & ~ (e21 = e24) & ~ (e21 = e23) & ~ (e21 = e22) & ~ (e20 = e24) & ~ (e20 = e23) & ~ (e20 = e22) & ~ (e20 = e21)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG052+1.p', ax10)).
fof(f9773, plain, ((e21 = e24) | (~ spl68_339 | ~ spl68_450)), inference(backward_demodulation, [], [f3329, f2656])).
fof(f2656, plain, ((e21 = h3(e13)) | ~ spl68_339), inference(avatar_component_clause, [], [f2655])).
fof(f2655, plain, (spl68_339 <=> (e21 = h3(e13))), introduced(avatar_definition, [new_symbols(naming, [spl68_339])])).
fof(f3329, plain, ((e24 = h3(e13)) | ~ spl68_450), inference(avatar_component_clause, [], [f3328])).
fof(f3328, plain, (spl68_450 <=> (e24 = h3(e13))), introduced(avatar_definition, [new_symbols(naming, [spl68_450])])).
fof(f9718, plain, (~ spl68_85 | ~ spl68_215 | ~ spl68_314 | spl68_396), inference(avatar_contradiction_clause, [], [f9717])).
fof(f9717, plain, ($false | (~ spl68_85 | ~ spl68_215 | ~ spl68_314 | spl68_396)), inference(subsumption_resolution, [], [f758, f9564])).
fof(f9564, plain, (~ (e24 = h5(e14)) | (~ spl68_85 | ~ spl68_215 | ~ spl68_314 | spl68_396)), inference(backward_demodulation, [], [f9322, f1243])).
fof(f9322, plain, (~ (e24 = h5(op1(e11, e13))) | (~ spl68_215 | ~ spl68_314 | spl68_396)), inference(forward_demodulation, [], [f8747, f1839])).
fof(f8747, plain, (~ (op2(e21, e23) = h5(op1(e11, e13))) | (~ spl68_314 | spl68_396)), inference(forward_demodulation, [], [f2961, f2525])).
fof(f2961, plain, (~ (h5(op1(e11, e13)) = op2(h5(e11), e23)) | spl68_396), inference(avatar_component_clause, [], [f2959])).
fof(f9695, plain, (~ spl68_2 | ~ spl68_4), inference(avatar_contradiction_clause, [], [f9694])).
fof(f9694, plain, ($false | (~ spl68_2 | ~ spl68_4)), inference(subsumption_resolution, [], [f9693, f540])).
fof(f9693, plain, ((e11 = e13) | (~ spl68_2 | ~ spl68_4)), inference(forward_demodulation, [], [f903, f895])).
fof(f903, plain, ((e13 = op1(e14, e14)) | ~ spl68_4), inference(avatar_component_clause, [], [f901])).
fof(f901, plain, (spl68_4 <=> (e13 = op1(e14, e14))), introduced(avatar_definition, [new_symbols(naming, [spl68_4])])).
fof(f9651, plain, (~ spl68_34 | ~ spl68_39), inference(avatar_split_clause, [], [f9646, f1048, f1027])).
fof(f1027, plain, (spl68_34 <=> (e13 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl68_34])])).
fof(f9646, plain, (~ (e13 = op1(e13, e13)) | ~ spl68_39), inference(backward_demodulation, [], [f422, f1050])).
fof(f1050, plain, ((e13 = op1(e13, e12)) | ~ spl68_39), inference(avatar_component_clause, [], [f1048])).
fof(f422, plain, ~ (op1(e13, e12) = op1(e13, e13)), inference(cnf_transformation, [], [f7])).
fof(f9568, plain, (~ spl68_35 | ~ spl68_85), inference(avatar_split_clause, [], [f9562, f1241, f1031])).
fof(f9562, plain, (~ (e14 = op1(e13, e13)) | ~ spl68_85), inference(backward_demodulation, [], [f370, f1243])).
fof(f370, plain, ~ (op1(e11, e13) = op1(e13, e13)), inference(cnf_transformation, [], [f7])).
fof(f9543, plain, (~ spl68_93 | ~ spl68_223 | ~ spl68_310 | ~ spl68_314 | spl68_378), inference(avatar_contradiction_clause, [], [f9542])).
fof(f9542, plain, ($false | (~ spl68_93 | ~ spl68_223 | ~ spl68_310 | ~ spl68_314 | spl68_378)), inference(subsumption_resolution, [], [f9535, f2503])).
fof(f9535, plain, (~ (e22 = h5(e12)) | (~ spl68_93 | ~ spl68_223 | ~ spl68_314 | spl68_378)), inference(backward_demodulation, [], [f9329, f1277])).
fof(f9329, plain, (~ (e22 = h5(op1(e11, e11))) | (~ spl68_223 | ~ spl68_314 | spl68_378)), inference(forward_demodulation, [], [f8570, f1873])).
fof(f1873, plain, ((e22 = op2(e21, e21)) | ~ spl68_223), inference(avatar_component_clause, [], [f1871])).
fof(f1871, plain, (spl68_223 <=> (e22 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl68_223])])).
fof(f8570, plain, (~ (op2(e21, e21) = h5(op1(e11, e11))) | (~ spl68_314 | spl68_378)), inference(forward_demodulation, [], [f2889, f2525])).
fof(f2889, plain, (~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | spl68_378), inference(avatar_component_clause, [], [f2887])).
fof(f2887, plain, (spl68_378 <=> (h5(op1(e11, e11)) = op2(h5(e11), h5(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl68_378])])).
fof(f9538, plain, (~ spl68_83 | ~ spl68_93), inference(avatar_split_clause, [], [f9531, f1275, f1233])).
fof(f9531, plain, (~ (e12 = op1(e11, e13)) | ~ spl68_93), inference(backward_demodulation, [], [f400, f1277])).
fof(f9475, plain, (~ spl68_95 | ~ spl68_120), inference(avatar_split_clause, [], [f9468, f1388, f1283])).
fof(f1388, plain, (spl68_120 <=> (e14 = op1(e10, e11))), introduced(avatar_definition, [new_symbols(naming, [spl68_120])])).
fof(f9468, plain, (~ (e14 = op1(e11, e11)) | ~ spl68_120), inference(backward_demodulation, [], [f345, f1390])).
fof(f1390, plain, ((e14 = op1(e10, e11)) | ~ spl68_120), inference(avatar_component_clause, [], [f1388])).
fof(f345, plain, ~ (op1(e10, e11) = op1(e11, e11)), inference(cnf_transformation, [], [f7])).
fof(f9341, plain, (~ spl68_107 | ~ spl68_158 | ~ spl68_337 | ~ spl68_341 | ~ spl68_437 | ~ spl68_450), inference(avatar_contradiction_clause, [], [f9340])).
fof(f9340, plain, ($false | (~ spl68_107 | ~ spl68_158 | ~ spl68_337 | ~ spl68_341 | ~ spl68_437 | ~ spl68_450)), inference(subsumption_resolution, [], [f9339, f549])).
fof(f549, plain, ~ (e21 = e22), inference(cnf_transformation, [], [f10])).
fof(f9339, plain, ((e21 = e22) | (~ spl68_107 | ~ spl68_158 | ~ spl68_337 | ~ spl68_341 | ~ spl68_437 | ~ spl68_450)), inference(forward_demodulation, [], [f9338, f2666])).
fof(f2666, plain, ((e21 = h3(e11)) | ~ spl68_341), inference(avatar_component_clause, [], [f2665])).
fof(f2665, plain, (spl68_341 <=> (e21 = h3(e11))), introduced(avatar_definition, [new_symbols(naming, [spl68_341])])).
fof(f9338, plain, ((e22 = h3(e11)) | (~ spl68_107 | ~ spl68_158 | ~ spl68_337 | ~ spl68_437 | ~ spl68_450)), inference(forward_demodulation, [], [f8846, f1336])).
fof(f8846, plain, ((e22 = h3(op1(e10, e13))) | (~ spl68_158 | ~ spl68_337 | ~ spl68_437 | ~ spl68_450)), inference(forward_demodulation, [], [f8837, f1600])).
fof(f8837, plain, ((op2(e23, e24) = h3(op1(e10, e13))) | (~ spl68_337 | ~ spl68_437 | ~ spl68_450)), inference(backward_demodulation, [], [f8727, f2645])).
fof(f2645, plain, ((e23 = h3(e10)) | ~ spl68_337), inference(avatar_component_clause, [], [f2644])).
fof(f2644, plain, (spl68_337 <=> (e23 = h3(e10))), introduced(avatar_definition, [new_symbols(naming, [spl68_337])])).
fof(f8727, plain, ((h3(op1(e10, e13)) = op2(h3(e10), e24)) | (~ spl68_437 | ~ spl68_450)), inference(backward_demodulation, [], [f3277, f3329])).
fof(f3277, plain, ((h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ spl68_437), inference(avatar_component_clause, [], [f3276])).
fof(f3276, plain, (spl68_437 <=> (h3(op1(e10, e13)) = op2(h3(e10), h3(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl68_437])])).
fof(f9280, plain, (~ spl68_27 | ~ spl68_28), inference(avatar_contradiction_clause, [], [f9279])).
fof(f9279, plain, ($false | (~ spl68_27 | ~ spl68_28)), inference(subsumption_resolution, [], [f9278, f539])).
fof(f539, plain, ~ (e11 = e12), inference(cnf_transformation, [], [f9])).
fof(f9278, plain, ((e11 = e12) | (~ spl68_27 | ~ spl68_28)), inference(backward_demodulation, [], [f1004, f1000])).
fof(f1000, plain, ((e11 = op1(e13, e14)) | ~ spl68_27), inference(avatar_component_clause, [], [f998])).
fof(f998, plain, (spl68_27 <=> (e11 = op1(e13, e14))), introduced(avatar_definition, [new_symbols(naming, [spl68_27])])).
fof(f9277, plain, (~ spl68_34 | ~ spl68_132 | spl68_339 | ~ spl68_449 | ~ spl68_450), inference(avatar_contradiction_clause, [], [f9276])).
fof(f9276, plain, ($false | (~ spl68_34 | ~ spl68_132 | spl68_339 | ~ spl68_449 | ~ spl68_450)), inference(subsumption_resolution, [], [f9272, f2657])).
fof(f2657, plain, (~ (e21 = h3(e13)) | spl68_339), inference(avatar_component_clause, [], [f2655])).
fof(f9272, plain, ((e21 = h3(e13)) | (~ spl68_34 | ~ spl68_132 | ~ spl68_449 | ~ spl68_450)), inference(backward_demodulation, [], [f8742, f1029])).
fof(f1029, plain, ((e13 = op1(e13, e13)) | ~ spl68_34), inference(avatar_component_clause, [], [f1027])).
fof(f8742, plain, ((e21 = h3(op1(e13, e13))) | (~ spl68_132 | ~ spl68_449 | ~ spl68_450)), inference(forward_demodulation, [], [f8731, f1491])).
fof(f8731, plain, ((op2(e24, e24) = h3(op1(e13, e13))) | (~ spl68_449 | ~ spl68_450)), inference(backward_demodulation, [], [f3325, f3329])).
fof(f3325, plain, ((h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ spl68_449), inference(avatar_component_clause, [], [f3324])).
fof(f3324, plain, (spl68_449 <=> (h3(op1(e13, e13)) = op2(h3(e13), h3(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl68_449])])).
fof(f9263, plain, (~ spl68_15 | ~ spl68_40), inference(avatar_split_clause, [], [f9258, f1052, f947])).
fof(f9258, plain, (~ (e14 = op1(e14, e12)) | ~ spl68_40), inference(backward_demodulation, [], [f364, f1054])).
fof(f9257, plain, (~ spl68_49 | ~ spl68_179 | spl68_391), inference(avatar_contradiction_clause, [], [f9256])).
fof(f9256, plain, ($false | (~ spl68_49 | ~ spl68_179 | spl68_391)), inference(subsumption_resolution, [], [f9249, f2472])).
fof(f9249, plain, (~ (e23 = h5(e13)) | (~ spl68_49 | ~ spl68_179 | spl68_391)), inference(backward_demodulation, [], [f8991, f1092])).
fof(f8991, plain, (~ (e23 = h5(op1(e13, e10))) | (~ spl68_179 | spl68_391)), inference(backward_demodulation, [], [f2941, f1688])).
fof(f1688, plain, ((e23 = op2(e23, e20)) | ~ spl68_179), inference(avatar_component_clause, [], [f1686])).
fof(f1686, plain, (spl68_179 <=> (e23 = op2(e23, e20))), introduced(avatar_definition, [new_symbols(naming, [spl68_179])])).
fof(f9251, plain, (~ spl68_39 | ~ spl68_49), inference(avatar_split_clause, [], [f9245, f1090, f1048])).
fof(f9245, plain, (~ (e13 = op1(e13, e12)) | ~ spl68_49), inference(backward_demodulation, [], [f416, f1092])).
fof(f416, plain, ~ (op1(e13, e10) = op1(e13, e12)), inference(cnf_transformation, [], [f7])).
fof(f9175, plain, (~ spl68_68 | ~ spl68_73), inference(avatar_split_clause, [], [f9167, f1191, f1170])).
fof(f1170, plain, (spl68_68 <=> (e12 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl68_68])])).
fof(f9167, plain, (~ (e12 = op1(e12, e11)) | ~ spl68_73), inference(backward_demodulation, [], [f405, f1193])).
fof(f405, plain, ~ (op1(e12, e10) = op1(e12, e11)), inference(cnf_transformation, [], [f7])).
fof(f9115, plain, (~ spl68_109 | ~ spl68_239 | spl68_399), inference(avatar_contradiction_clause, [], [f9114])).
fof(f9114, plain, ($false | (~ spl68_109 | ~ spl68_239 | spl68_399)), inference(subsumption_resolution, [], [f9105, f2472])).
fof(f9105, plain, (~ (e23 = h5(e13)) | (~ spl68_109 | ~ spl68_239 | spl68_399)), inference(backward_demodulation, [], [f8943, f1344])).
fof(f8943, plain, (~ (e23 = h5(op1(e10, e13))) | (~ spl68_239 | spl68_399)), inference(backward_demodulation, [], [f2973, f1940])).
fof(f1940, plain, ((e23 = op2(e20, e23)) | ~ spl68_239), inference(avatar_component_clause, [], [f1938])).
fof(f1938, plain, (spl68_239 <=> (e23 = op2(e20, e23))), introduced(avatar_definition, [new_symbols(naming, [spl68_239])])).
fof(f9108, plain, (~ spl68_34 | ~ spl68_109), inference(avatar_split_clause, [], [f9100, f1342, f1027])).
fof(f9100, plain, (~ (e13 = op1(e13, e13)) | ~ spl68_109), inference(backward_demodulation, [], [f367, f1344])).
fof(f367, plain, ~ (op1(e10, e13) = op1(e13, e13)), inference(cnf_transformation, [], [f7])).
fof(f9095, plain, (~ spl68_13 | ~ spl68_113), inference(avatar_split_clause, [], [f9086, f1359, f939])).
fof(f939, plain, (spl68_13 <=> (e12 = op1(e14, e12))), introduced(avatar_definition, [new_symbols(naming, [spl68_13])])).
fof(f9086, plain, (~ (e12 = op1(e14, e12)) | ~ spl68_113), inference(backward_demodulation, [], [f358, f1361])).
fof(f358, plain, ~ (op1(e10, e12) = op1(e14, e12)), inference(cnf_transformation, [], [f7])).
fof(f9059, plain, (~ spl68_106 | ~ spl68_121), inference(avatar_split_clause, [], [f9048, f1393, f1330])).
fof(f1330, plain, (spl68_106 <=> (e10 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl68_106])])).
fof(f9048, plain, (~ (e10 = op1(e10, e13)) | ~ spl68_121), inference(backward_demodulation, [], [f387, f1395])).
fof(f387, plain, ~ (op1(e10, e10) = op1(e10, e13)), inference(cnf_transformation, [], [f7])).
fof(f9039, plain, (spl68_49 | ~ spl68_126), inference(avatar_split_clause, [], [f9025, f1414, f1090])).
fof(f9025, plain, ((e13 = op1(e13, e10)) | ~ spl68_126), inference(backward_demodulation, [], [f195, f1416])).
fof(f8993, plain, (~ spl68_169 | ~ spl68_179), inference(avatar_split_clause, [], [f8989, f1686, f1644])).
fof(f8989, plain, (~ (e23 = op2(e23, e22)) | ~ spl68_179), inference(backward_demodulation, [], [f516, f1688])).
fof(f516, plain, ~ (op2(e23, e20) = op2(e23, e22)), inference(cnf_transformation, [], [f8])).
fof(f8959, plain, (~ spl68_177 | ~ spl68_227), inference(avatar_split_clause, [], [f8952, f1888, f1678])).
fof(f8952, plain, (~ (e21 = op2(e23, e20)) | ~ spl68_227), inference(backward_demodulation, [], [f440, f1890])).
fof(f440, plain, ~ (op2(e21, e20) = op2(e23, e20)), inference(cnf_transformation, [], [f8])).
fof(f8950, plain, (~ spl68_185 | ~ spl68_235), inference(avatar_split_clause, [], [f8947, f1921, f1711])).
fof(f8947, plain, (~ (e24 = op2(e22, e24)) | ~ spl68_235), inference(backward_demodulation, [], [f476, f1923])).
fof(f476, plain, ~ (op2(e20, e24) = op2(e22, e24)), inference(cnf_transformation, [], [f8])).
fof(f8916, plain, (~ spl68_241 | ~ spl68_251), inference(avatar_split_clause, [], [f8905, f1989, f1947])).
fof(f8905, plain, (~ (e20 = op2(e20, e22)) | ~ spl68_251), inference(backward_demodulation, [], [f2434, f8897])).
fof(f2434, plain, ~ (op2(e20, e22) = h1(e11)), inference(backward_demodulation, [], [f486, f740])).
fof(f486, plain, ~ (op2(e20, e20) = op2(e20, e22)), inference(cnf_transformation, [], [f8])).
fof(f8915, plain, (~ spl68_201 | ~ spl68_251), inference(avatar_split_clause, [], [f8901, f1989, f1779])).
fof(f8901, plain, (~ (e20 = op2(e22, e20)) | ~ spl68_251), inference(backward_demodulation, [], [f2430, f8897])).
fof(f2430, plain, ~ (op2(e22, e20) = h1(e11)), inference(backward_demodulation, [], [f436, f740])).
fof(f436, plain, ~ (op2(e20, e20) = op2(e22, e20)), inference(cnf_transformation, [], [f8])).
fof(f8894, plain, (spl68_155 | ~ spl68_256), inference(avatar_split_clause, [], [f8880, f2010, f1585])).
fof(f2010, plain, (spl68_256 <=> (e20 = unit2)), introduced(avatar_definition, [new_symbols(naming, [spl68_256])])).
fof(f8880, plain, ((e24 = op2(e24, e20)) | ~ spl68_256), inference(backward_demodulation, [], [f283, f2012])).
fof(f2012, plain, ((e20 = unit2) | ~ spl68_256), inference(avatar_component_clause, [], [f2010])).
fof(f8893, plain, (spl68_235 | ~ spl68_256), inference(avatar_split_clause, [], [f8879, f2010, f1921])).
fof(f8879, plain, ((e24 = op2(e20, e24)) | ~ spl68_256), inference(backward_demodulation, [], [f282, f2012])).
fof(f8892, plain, (spl68_179 | ~ spl68_256), inference(avatar_split_clause, [], [f8878, f2010, f1686])).
fof(f8878, plain, ((e23 = op2(e23, e20)) | ~ spl68_256), inference(backward_demodulation, [], [f281, f2012])).
fof(f8891, plain, (spl68_239 | ~ spl68_256), inference(avatar_split_clause, [], [f8877, f2010, f1938])).
fof(f8877, plain, ((e23 = op2(e20, e23)) | ~ spl68_256), inference(backward_demodulation, [], [f280, f2012])).
fof(f8890, plain, (spl68_203 | ~ spl68_256), inference(avatar_split_clause, [], [f8876, f2010, f1787])).
fof(f8876, plain, ((e22 = op2(e22, e20)) | ~ spl68_256), inference(backward_demodulation, [], [f279, f2012])).
fof(f8889, plain, (spl68_243 | ~ spl68_256), inference(avatar_split_clause, [], [f8875, f2010, f1955])).
fof(f8875, plain, ((e22 = op2(e20, e22)) | ~ spl68_256), inference(backward_demodulation, [], [f278, f2012])).
fof(f278, plain, (e22 = op2(unit2, e22)), inference(cnf_transformation, [], [f5])).
fof(f8888, plain, (spl68_227 | ~ spl68_256), inference(avatar_split_clause, [], [f8874, f2010, f1888])).
fof(f8874, plain, ((e21 = op2(e21, e20)) | ~ spl68_256), inference(backward_demodulation, [], [f277, f2012])).
fof(f8887, plain, (spl68_247 | ~ spl68_256), inference(avatar_split_clause, [], [f8873, f2010, f1972])).
fof(f8873, plain, ((e21 = op2(e20, e21)) | ~ spl68_256), inference(backward_demodulation, [], [f276, f2012])).
fof(f8886, plain, (spl68_251 | ~ spl68_256), inference(avatar_split_clause, [], [f8872, f2010, f1989])).
fof(f8872, plain, ((e20 = op2(e20, e20)) | ~ spl68_256), inference(backward_demodulation, [], [f275, f2012])).
fof(f8624, plain, (~ spl68_51 | ~ spl68_76), inference(avatar_split_clause, [], [f4829, f1204, f1099])).
fof(f1099, plain, (spl68_51 <=> (e10 = op1(e12, e14))), introduced(avatar_definition, [new_symbols(naming, [spl68_51])])).
fof(f4829, plain, (~ (e10 = op1(e12, e14)) | ~ spl68_76), inference(forward_demodulation, [], [f379, f1206])).
fof(f379, plain, ~ (op1(e11, e14) = op1(e12, e14)), inference(cnf_transformation, [], [f7])).
fof(f8573, plain, (~ spl68_225 | ~ spl68_95 | ~ spl68_314 | spl68_378), inference(avatar_split_clause, [], [f8572, f2887, f2524, f1283, f1879])).
fof(f8572, plain, (~ (e24 = op2(e21, e21)) | (~ spl68_95 | ~ spl68_314 | spl68_378)), inference(forward_demodulation, [], [f8571, f758])).
fof(f8571, plain, (~ (op2(e21, e21) = h5(e14)) | (~ spl68_95 | ~ spl68_314 | spl68_378)), inference(forward_demodulation, [], [f8570, f1285])).
fof(f8536, plain, (~ spl68_136 | ~ spl68_331), inference(avatar_split_clause, [], [f8535, f2614, f1506])).
fof(f2614, plain, (spl68_331 <=> (e20 = h4(e11))), introduced(avatar_definition, [new_symbols(naming, [spl68_331])])).
fof(f8535, plain, (~ (e20 = op2(e24, e23)) | ~ spl68_331), inference(forward_demodulation, [], [f2465, f2615])).
fof(f2615, plain, ((e20 = h4(e11)) | ~ spl68_331), inference(avatar_component_clause, [], [f2614])).
fof(f2465, plain, ~ (op2(e24, e23) = h4(e11)), inference(backward_demodulation, [], [f474, f755])).
fof(f474, plain, ~ (op2(e23, e23) = op2(e24, e23)), inference(cnf_transformation, [], [f8])).
fof(f8533, plain, (~ spl68_366 | ~ spl68_234), inference(avatar_split_clause, [], [f7239, f1917, f2791])).
fof(f2791, plain, (spl68_366 <=> (e23 = h1(e11))), introduced(avatar_definition, [new_symbols(naming, [spl68_366])])).
fof(f7239, plain, (~ (e23 = h1(e11)) | ~ spl68_234), inference(backward_demodulation, [], [f2436, f1919])).
fof(f2436, plain, ~ (op2(e20, e24) = h1(e11)), inference(backward_demodulation, [], [f488, f740])).
fof(f488, plain, ~ (op2(e20, e20) = op2(e20, e24)), inference(cnf_transformation, [], [f8])).
fof(f8531, plain, (~ spl68_361 | ~ spl68_206), inference(avatar_split_clause, [], [f5293, f1800, f2766])).
fof(f2766, plain, (spl68_361 <=> (e20 = h2(e11))), introduced(avatar_definition, [new_symbols(naming, [spl68_361])])).
fof(f5293, plain, (~ (e20 = h2(e11)) | ~ spl68_206), inference(forward_demodulation, [], [f2447, f1802])).
fof(f2447, plain, ~ (op2(e21, e24) = h2(e11)), inference(backward_demodulation, [], [f501, f745])).
fof(f501, plain, ~ (op2(e21, e21) = op2(e21, e24)), inference(cnf_transformation, [], [f8])).
fof(f8503, plain, (~ spl68_146 | ~ spl68_149), inference(avatar_contradiction_clause, [], [f8502])).
fof(f8502, plain, ($false | (~ spl68_146 | ~ spl68_149)), inference(subsumption_resolution, [], [f8500, f547])).
fof(f547, plain, ~ (e20 = e23), inference(cnf_transformation, [], [f10])).
fof(f8500, plain, ((e20 = e23) | (~ spl68_146 | ~ spl68_149)), inference(backward_demodulation, [], [f1562, f1550])).
fof(f1550, plain, ((e20 = op2(e24, e21)) | ~ spl68_146), inference(avatar_component_clause, [], [f1548])).
fof(f1548, plain, (spl68_146 <=> (e20 = op2(e24, e21))), introduced(avatar_definition, [new_symbols(naming, [spl68_146])])).
fof(f8486, plain, (~ spl68_136 | ~ spl68_99 | ~ spl68_352 | ~ spl68_359 | spl68_467 | ~ spl68_490), inference(avatar_split_clause, [], [f8485, f3580, f3458, f2756, f2720, f1300, f1506])).
fof(f3458, plain, (spl68_467 <=> (h2(op1(e11, e10)) = op2(h2(e11), h2(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl68_467])])).
fof(f8485, plain, (~ (e20 = op2(e24, e23)) | (~ spl68_99 | ~ spl68_352 | ~ spl68_359 | spl68_467 | ~ spl68_490)), inference(forward_demodulation, [], [f8484, f2757])).
fof(f8484, plain, (~ (op2(e24, e23) = h2(e13)) | (~ spl68_99 | ~ spl68_352 | spl68_467 | ~ spl68_490)), inference(forward_demodulation, [], [f8483, f1302])).
fof(f8483, plain, (~ (op2(e24, e23) = h2(op1(e11, e10))) | (~ spl68_352 | spl68_467 | ~ spl68_490)), inference(forward_demodulation, [], [f8482, f3581])).
fof(f8482, plain, (~ (h2(op1(e11, e10)) = op2(h2(e11), e23)) | (~ spl68_352 | spl68_467)), inference(forward_demodulation, [], [f3460, f2721])).
fof(f3460, plain, (~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | spl68_467), inference(avatar_component_clause, [], [f3458])).
fof(f8475, plain, (~ spl68_172 | ~ spl68_197), inference(avatar_split_clause, [], [f8474, f1762, f1657])).
fof(f1657, plain, (spl68_172 <=> (e21 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl68_172])])).
fof(f8474, plain, (~ (e21 = op2(e23, e21)) | ~ spl68_197), inference(forward_demodulation, [], [f452, f1764])).
fof(f452, plain, ~ (op2(e22, e21) = op2(e23, e21)), inference(cnf_transformation, [], [f8])).
fof(f8458, plain, (~ spl68_138 | ~ spl68_213), inference(avatar_split_clause, [], [f8457, f1829, f1514])).
fof(f8457, plain, (~ (e22 = op2(e24, e23)) | ~ spl68_213), inference(forward_demodulation, [], [f471, f1831])).
fof(f471, plain, ~ (op2(e21, e23) = op2(e24, e23)), inference(cnf_transformation, [], [f8])).
fof(f8452, plain, (~ spl68_67 | ~ spl68_185 | ~ spl68_355 | spl68_472 | ~ spl68_490), inference(avatar_contradiction_clause, [], [f8451])).
fof(f8451, plain, ($false | (~ spl68_67 | ~ spl68_185 | ~ spl68_355 | spl68_472 | ~ spl68_490)), inference(subsumption_resolution, [], [f8450, f1713])).
fof(f8450, plain, (~ (e24 = op2(e22, e24)) | (~ spl68_67 | ~ spl68_355 | spl68_472 | ~ spl68_490)), inference(forward_demodulation, [], [f8449, f3581])).
fof(f8449, plain, (~ (op2(e22, e24) = h2(e11)) | (~ spl68_67 | ~ spl68_355 | spl68_472 | ~ spl68_490)), inference(forward_demodulation, [], [f8448, f1168])).
fof(f8448, plain, (~ (op2(e22, e24) = h2(op1(e12, e11))) | (~ spl68_355 | spl68_472 | ~ spl68_490)), inference(forward_demodulation, [], [f8447, f2736])).
fof(f2736, plain, ((e22 = h2(e12)) | ~ spl68_355), inference(avatar_component_clause, [], [f2735])).
fof(f2735, plain, (spl68_355 <=> (e22 = h2(e12))), introduced(avatar_definition, [new_symbols(naming, [spl68_355])])).
fof(f8447, plain, (~ (h2(op1(e12, e11)) = op2(h2(e12), e24)) | (spl68_472 | ~ spl68_490)), inference(forward_demodulation, [], [f3480, f3581])).
fof(f3480, plain, (~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | spl68_472), inference(avatar_component_clause, [], [f3478])).
fof(f3478, plain, (spl68_472 <=> (h2(op1(e12, e11)) = op2(h2(e12), h2(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl68_472])])).
fof(f8442, plain, (~ spl68_59 | ~ spl68_201 | ~ spl68_355 | ~ spl68_359 | spl68_474), inference(avatar_contradiction_clause, [], [f8441])).
fof(f8441, plain, ($false | (~ spl68_59 | ~ spl68_201 | ~ spl68_355 | ~ spl68_359 | spl68_474)), inference(subsumption_resolution, [], [f8440, f1781])).
fof(f8440, plain, (~ (e20 = op2(e22, e20)) | (~ spl68_59 | ~ spl68_355 | ~ spl68_359 | spl68_474)), inference(forward_demodulation, [], [f8439, f2757])).
fof(f8439, plain, (~ (op2(e22, e20) = h2(e13)) | (~ spl68_59 | ~ spl68_355 | ~ spl68_359 | spl68_474)), inference(forward_demodulation, [], [f8438, f1134])).
fof(f1134, plain, ((e13 = op1(e12, e13)) | ~ spl68_59), inference(avatar_component_clause, [], [f1132])).
fof(f8438, plain, (~ (op2(e22, e20) = h2(op1(e12, e13))) | (~ spl68_355 | ~ spl68_359 | spl68_474)), inference(forward_demodulation, [], [f8437, f2736])).
fof(f8437, plain, (~ (h2(op1(e12, e13)) = op2(h2(e12), e20)) | (~ spl68_359 | spl68_474)), inference(forward_demodulation, [], [f3488, f2757])).
fof(f3488, plain, (~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | spl68_474), inference(avatar_component_clause, [], [f3486])).
fof(f3486, plain, (spl68_474 <=> (h2(op1(e12, e13)) = op2(h2(e12), h2(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl68_474])])).
fof(f8429, plain, (~ spl68_15 | ~ spl68_217 | ~ spl68_355 | spl68_482), inference(avatar_contradiction_clause, [], [f8428])).
fof(f8428, plain, ($false | (~ spl68_15 | ~ spl68_217 | ~ spl68_355 | spl68_482)), inference(subsumption_resolution, [], [f8427, f1848])).
fof(f8427, plain, (~ (e21 = op2(e21, e22)) | (~ spl68_15 | ~ spl68_355 | spl68_482)), inference(forward_demodulation, [], [f8426, f743])).
fof(f8426, plain, (~ (op2(e21, e22) = h2(e14)) | (~ spl68_15 | ~ spl68_355 | spl68_482)), inference(forward_demodulation, [], [f8425, f949])).
fof(f949, plain, ((e14 = op1(e14, e12)) | ~ spl68_15), inference(avatar_component_clause, [], [f947])).
fof(f8425, plain, (~ (op2(e21, e22) = h2(op1(e14, e12))) | (~ spl68_355 | spl68_482)), inference(forward_demodulation, [], [f3520, f2736])).
fof(f3520, plain, (~ (h2(op1(e14, e12)) = op2(e21, h2(e12))) | spl68_482), inference(avatar_component_clause, [], [f3518])).
fof(f3518, plain, (spl68_482 <=> (h2(op1(e14, e12)) = op2(e21, h2(e12)))), introduced(avatar_definition, [new_symbols(naming, [spl68_482])])).
fof(f8414, plain, (~ spl68_111 | ~ spl68_169 | ~ spl68_352 | ~ spl68_355 | spl68_465), inference(avatar_contradiction_clause, [], [f8413])).
fof(f8413, plain, ($false | (~ spl68_111 | ~ spl68_169 | ~ spl68_352 | ~ spl68_355 | spl68_465)), inference(subsumption_resolution, [], [f8412, f1646])).
fof(f8412, plain, (~ (e23 = op2(e23, e22)) | (~ spl68_111 | ~ spl68_352 | ~ spl68_355 | spl68_465)), inference(forward_demodulation, [], [f8411, f2721])).
fof(f8411, plain, (~ (op2(e23, e22) = h2(e10)) | (~ spl68_111 | ~ spl68_352 | ~ spl68_355 | spl68_465)), inference(forward_demodulation, [], [f8410, f1353])).
fof(f1353, plain, ((e10 = op1(e10, e12)) | ~ spl68_111), inference(avatar_component_clause, [], [f1351])).
fof(f8410, plain, (~ (op2(e23, e22) = h2(op1(e10, e12))) | (~ spl68_352 | ~ spl68_355 | spl68_465)), inference(forward_demodulation, [], [f8409, f2721])).
fof(f8409, plain, (~ (h2(op1(e10, e12)) = op2(h2(e10), e22)) | (~ spl68_355 | spl68_465)), inference(forward_demodulation, [], [f3452, f2736])).
fof(f3452, plain, (~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | spl68_465), inference(avatar_component_clause, [], [f3450])).
fof(f3450, plain, (spl68_465 <=> (h2(op1(e10, e12)) = op2(h2(e10), h2(e12)))), introduced(avatar_definition, [new_symbols(naming, [spl68_465])])).
fof(f8404, plain, (~ spl68_83 | ~ spl68_153 | ~ spl68_355 | ~ spl68_359 | spl68_470 | ~ spl68_490), inference(avatar_contradiction_clause, [], [f8403])).
fof(f8403, plain, ($false | (~ spl68_83 | ~ spl68_153 | ~ spl68_355 | ~ spl68_359 | spl68_470 | ~ spl68_490)), inference(subsumption_resolution, [], [f8402, f1579])).
fof(f8402, plain, (~ (e22 = op2(e24, e20)) | (~ spl68_83 | ~ spl68_355 | ~ spl68_359 | spl68_470 | ~ spl68_490)), inference(forward_demodulation, [], [f8401, f2736])).
fof(f8401, plain, (~ (op2(e24, e20) = h2(e12)) | (~ spl68_83 | ~ spl68_359 | spl68_470 | ~ spl68_490)), inference(forward_demodulation, [], [f8400, f1235])).
fof(f8400, plain, (~ (op2(e24, e20) = h2(op1(e11, e13))) | (~ spl68_359 | spl68_470 | ~ spl68_490)), inference(forward_demodulation, [], [f8399, f3581])).
fof(f8399, plain, (~ (h2(op1(e11, e13)) = op2(h2(e11), e20)) | (~ spl68_359 | spl68_470)), inference(forward_demodulation, [], [f3472, f2757])).
fof(f3472, plain, (~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | spl68_470), inference(avatar_component_clause, [], [f3470])).
fof(f3470, plain, (spl68_470 <=> (h2(op1(e11, e13)) = op2(h2(e11), h2(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl68_470])])).
fof(f8392, plain, (~ spl68_41 | ~ spl68_234 | ~ spl68_352 | ~ spl68_359 | spl68_476 | ~ spl68_490), inference(avatar_contradiction_clause, [], [f8391])).
fof(f8391, plain, ($false | (~ spl68_41 | ~ spl68_234 | ~ spl68_352 | ~ spl68_359 | spl68_476 | ~ spl68_490)), inference(subsumption_resolution, [], [f8390, f1919])).
fof(f8390, plain, (~ (e23 = op2(e20, e24)) | (~ spl68_41 | ~ spl68_352 | ~ spl68_359 | spl68_476 | ~ spl68_490)), inference(forward_demodulation, [], [f8389, f2721])).
fof(f8389, plain, (~ (op2(e20, e24) = h2(e10)) | (~ spl68_41 | ~ spl68_359 | spl68_476 | ~ spl68_490)), inference(forward_demodulation, [], [f8388, f1059])).
fof(f8388, plain, (~ (op2(e20, e24) = h2(op1(e13, e11))) | (~ spl68_359 | spl68_476 | ~ spl68_490)), inference(forward_demodulation, [], [f8387, f2757])).
fof(f8387, plain, (~ (h2(op1(e13, e11)) = op2(h2(e13), e24)) | (spl68_476 | ~ spl68_490)), inference(forward_demodulation, [], [f3496, f3581])).
fof(f3496, plain, (~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | spl68_476), inference(avatar_component_clause, [], [f3494])).
fof(f3494, plain, (spl68_476 <=> (h2(op1(e13, e11)) = op2(h2(e13), h2(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl68_476])])).
fof(f8378, plain, (~ spl68_39 | ~ spl68_241 | ~ spl68_355 | ~ spl68_359 | spl68_477), inference(avatar_contradiction_clause, [], [f8377])).
fof(f8377, plain, ($false | (~ spl68_39 | ~ spl68_241 | ~ spl68_355 | ~ spl68_359 | spl68_477)), inference(subsumption_resolution, [], [f8376, f1949])).
fof(f8376, plain, (~ (e20 = op2(e20, e22)) | (~ spl68_39 | ~ spl68_355 | ~ spl68_359 | spl68_477)), inference(forward_demodulation, [], [f8375, f2757])).
fof(f8375, plain, (~ (op2(e20, e22) = h2(e13)) | (~ spl68_39 | ~ spl68_355 | ~ spl68_359 | spl68_477)), inference(forward_demodulation, [], [f8374, f1050])).
fof(f8374, plain, (~ (op2(e20, e22) = h2(op1(e13, e12))) | (~ spl68_355 | ~ spl68_359 | spl68_477)), inference(forward_demodulation, [], [f8373, f2757])).
fof(f8373, plain, (~ (h2(op1(e13, e12)) = op2(h2(e13), e22)) | (~ spl68_355 | spl68_477)), inference(forward_demodulation, [], [f3500, f2736])).
fof(f3500, plain, (~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | spl68_477), inference(avatar_component_clause, [], [f3498])).
fof(f3498, plain, (spl68_477 <=> (h2(op1(e13, e12)) = op2(h2(e13), h2(e12)))), introduced(avatar_definition, [new_symbols(naming, [spl68_477])])).
fof(f8372, plain, (~ spl68_118 | ~ spl68_158 | ~ spl68_352 | ~ spl68_355 | spl68_464 | ~ spl68_490), inference(avatar_contradiction_clause, [], [f8371])).
fof(f8371, plain, ($false | (~ spl68_118 | ~ spl68_158 | ~ spl68_352 | ~ spl68_355 | spl68_464 | ~ spl68_490)), inference(subsumption_resolution, [], [f8370, f1600])).
fof(f8370, plain, (~ (e22 = op2(e23, e24)) | (~ spl68_118 | ~ spl68_352 | ~ spl68_355 | spl68_464 | ~ spl68_490)), inference(forward_demodulation, [], [f8369, f2736])).
fof(f8369, plain, (~ (op2(e23, e24) = h2(e12)) | (~ spl68_118 | ~ spl68_352 | spl68_464 | ~ spl68_490)), inference(forward_demodulation, [], [f8368, f1382])).
fof(f8368, plain, (~ (op2(e23, e24) = h2(op1(e10, e11))) | (~ spl68_352 | spl68_464 | ~ spl68_490)), inference(forward_demodulation, [], [f8367, f2721])).
fof(f8367, plain, (~ (h2(op1(e10, e11)) = op2(h2(e10), e24)) | (spl68_464 | ~ spl68_490)), inference(forward_demodulation, [], [f3448, f3581])).
fof(f3448, plain, (~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | spl68_464), inference(avatar_component_clause, [], [f3446])).
fof(f3446, plain, (spl68_464 <=> (h2(op1(e10, e11)) = op2(h2(e10), h2(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl68_464])])).
fof(f8359, plain, (~ spl68_6 | ~ spl68_229 | ~ spl68_352 | ~ spl68_359 | spl68_481), inference(avatar_contradiction_clause, [], [f8358])).
fof(f8358, plain, ($false | (~ spl68_6 | ~ spl68_229 | ~ spl68_352 | ~ spl68_359 | spl68_481)), inference(subsumption_resolution, [], [f8357, f1898])).
fof(f8357, plain, (~ (e23 = op2(e21, e20)) | (~ spl68_6 | ~ spl68_352 | ~ spl68_359 | spl68_481)), inference(forward_demodulation, [], [f8356, f2721])).
fof(f8356, plain, (~ (op2(e21, e20) = h2(e10)) | (~ spl68_6 | ~ spl68_359 | spl68_481)), inference(forward_demodulation, [], [f8355, f912])).
fof(f912, plain, ((e10 = op1(e14, e13)) | ~ spl68_6), inference(avatar_component_clause, [], [f910])).
fof(f8355, plain, (~ (op2(e21, e20) = h2(op1(e14, e13))) | (~ spl68_359 | spl68_481)), inference(forward_demodulation, [], [f3516, f2757])).
fof(f3516, plain, (~ (h2(op1(e14, e13)) = op2(e21, h2(e13))) | spl68_481), inference(avatar_component_clause, [], [f3514])).
fof(f3514, plain, (spl68_481 <=> (h2(op1(e14, e13)) = op2(e21, h2(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl68_481])])).
fof(f8354, plain, (~ spl68_87 | ~ spl68_145 | ~ spl68_355 | spl68_469 | ~ spl68_490), inference(avatar_contradiction_clause, [], [f8353])).
fof(f8353, plain, ($false | (~ spl68_87 | ~ spl68_145 | ~ spl68_355 | spl68_469 | ~ spl68_490)), inference(subsumption_resolution, [], [f8352, f1545])).
fof(f8352, plain, (~ (e24 = op2(e24, e22)) | (~ spl68_87 | ~ spl68_355 | spl68_469 | ~ spl68_490)), inference(forward_demodulation, [], [f8351, f3581])).
fof(f8351, plain, (~ (op2(e24, e22) = h2(e11)) | (~ spl68_87 | ~ spl68_355 | spl68_469 | ~ spl68_490)), inference(forward_demodulation, [], [f8350, f1252])).
fof(f8350, plain, (~ (op2(e24, e22) = h2(op1(e11, e12))) | (~ spl68_355 | spl68_469 | ~ spl68_490)), inference(forward_demodulation, [], [f8349, f3581])).
fof(f8349, plain, (~ (h2(op1(e11, e12)) = op2(h2(e11), e22)) | (~ spl68_355 | spl68_469)), inference(forward_demodulation, [], [f3468, f2736])).
fof(f3468, plain, (~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | spl68_469), inference(avatar_component_clause, [], [f3466])).
fof(f3466, plain, (spl68_469 <=> (h2(op1(e11, e12)) = op2(h2(e11), h2(e12)))), introduced(avatar_definition, [new_symbols(naming, [spl68_469])])).
fof(f8342, plain, (~ spl68_63 | ~ spl68_193 | ~ spl68_355 | spl68_473), inference(avatar_contradiction_clause, [], [f8341])).
fof(f8341, plain, ($false | (~ spl68_63 | ~ spl68_193 | ~ spl68_355 | spl68_473)), inference(subsumption_resolution, [], [f8340, f1747])).
fof(f8340, plain, (~ (e22 = op2(e22, e22)) | (~ spl68_63 | ~ spl68_355 | spl68_473)), inference(forward_demodulation, [], [f8339, f2736])).
fof(f8339, plain, (~ (op2(e22, e22) = h2(e12)) | (~ spl68_63 | ~ spl68_355 | spl68_473)), inference(forward_demodulation, [], [f8338, f1151])).
fof(f1151, plain, ((e12 = op1(e12, e12)) | ~ spl68_63), inference(avatar_component_clause, [], [f1149])).
fof(f8338, plain, (~ (op2(e22, e22) = h2(op1(e12, e12))) | (~ spl68_355 | spl68_473)), inference(forward_demodulation, [], [f3484, f2736])).
fof(f3484, plain, (~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | spl68_473), inference(avatar_component_clause, [], [f3482])).
fof(f3482, plain, (spl68_473 <=> (h2(op1(e12, e12)) = op2(h2(e12), h2(e12)))), introduced(avatar_definition, [new_symbols(naming, [spl68_473])])).
fof(f8337, plain, (~ spl68_95 | ~ spl68_132 | spl68_468 | ~ spl68_490), inference(avatar_contradiction_clause, [], [f8336])).
fof(f8336, plain, ($false | (~ spl68_95 | ~ spl68_132 | spl68_468 | ~ spl68_490)), inference(subsumption_resolution, [], [f8335, f1491])).
fof(f8335, plain, (~ (e21 = op2(e24, e24)) | (~ spl68_95 | spl68_468 | ~ spl68_490)), inference(forward_demodulation, [], [f8334, f743])).
fof(f8334, plain, (~ (op2(e24, e24) = h2(e14)) | (~ spl68_95 | spl68_468 | ~ spl68_490)), inference(forward_demodulation, [], [f8333, f1285])).
fof(f8333, plain, (~ (op2(e24, e24) = h2(op1(e11, e11))) | (spl68_468 | ~ spl68_490)), inference(forward_demodulation, [], [f3464, f3581])).
fof(f3464, plain, (~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | spl68_468), inference(avatar_component_clause, [], [f3462])).
fof(f3462, plain, (spl68_468 <=> (h2(op1(e11, e11)) = op2(h2(e11), h2(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl68_468])])).
fof(f8324, plain, (~ spl68_23 | ~ spl68_213 | ~ spl68_352 | ~ spl68_355 | spl68_484), inference(avatar_contradiction_clause, [], [f8323])).
fof(f8323, plain, ($false | (~ spl68_23 | ~ spl68_213 | ~ spl68_352 | ~ spl68_355 | spl68_484)), inference(subsumption_resolution, [], [f8322, f1831])).
fof(f8322, plain, (~ (e22 = op2(e21, e23)) | (~ spl68_23 | ~ spl68_352 | ~ spl68_355 | spl68_484)), inference(forward_demodulation, [], [f8321, f2736])).
fof(f8321, plain, (~ (op2(e21, e23) = h2(e12)) | (~ spl68_23 | ~ spl68_352 | spl68_484)), inference(forward_demodulation, [], [f8320, f983])).
fof(f983, plain, ((e12 = op1(e14, e10)) | ~ spl68_23), inference(avatar_component_clause, [], [f981])).
fof(f981, plain, (spl68_23 <=> (e12 = op1(e14, e10))), introduced(avatar_definition, [new_symbols(naming, [spl68_23])])).
fof(f8320, plain, (~ (op2(e21, e23) = h2(op1(e14, e10))) | (~ spl68_352 | spl68_484)), inference(forward_demodulation, [], [f3528, f2721])).
fof(f3528, plain, (~ (h2(op1(e14, e10)) = op2(e21, h2(e10))) | spl68_484), inference(avatar_component_clause, [], [f3526])).
fof(f3526, plain, (spl68_484 <=> (h2(op1(e14, e10)) = op2(e21, h2(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl68_484])])).
fof(f8319, plain, (spl68_375 | ~ spl68_177 | ~ spl68_364), inference(avatar_split_clause, [], [f8318, f2781, f1678, f2836])).
fof(f8318, plain, ((e21 = h1(e12)) | (~ spl68_177 | ~ spl68_364)), inference(forward_demodulation, [], [f8316, f1680])).
fof(f1680, plain, ((e21 = op2(e23, e20)) | ~ spl68_177), inference(avatar_component_clause, [], [f1678])).
fof(f8316, plain, ((op2(e23, e20) = h1(e12)) | ~ spl68_364), inference(backward_demodulation, [], [f2428, f2782])).
fof(f2428, plain, (h1(e12) = op2(h1(e13), e20)), inference(forward_demodulation, [], [f741, f742])).
fof(f741, plain, (h1(e12) = op2(op2(e20, op2(e20, e20)), e20)), inference(cnf_transformation, [], [f16])).
fof(f8308, plain, (~ spl68_71 | ~ spl68_189 | ~ spl68_352 | ~ spl68_355 | spl68_471), inference(avatar_contradiction_clause, [], [f8307])).
fof(f8307, plain, ($false | (~ spl68_71 | ~ spl68_189 | ~ spl68_352 | ~ spl68_355 | spl68_471)), inference(subsumption_resolution, [], [f8306, f1730])).
fof(f8306, plain, (~ (e23 = op2(e22, e23)) | (~ spl68_71 | ~ spl68_352 | ~ spl68_355 | spl68_471)), inference(forward_demodulation, [], [f8305, f2721])).
fof(f8305, plain, (~ (op2(e22, e23) = h2(e10)) | (~ spl68_71 | ~ spl68_352 | ~ spl68_355 | spl68_471)), inference(forward_demodulation, [], [f8304, f1185])).
fof(f8304, plain, (~ (op2(e22, e23) = h2(op1(e12, e10))) | (~ spl68_352 | ~ spl68_355 | spl68_471)), inference(forward_demodulation, [], [f8303, f2736])).
fof(f8303, plain, (~ (h2(op1(e12, e10)) = op2(h2(e12), e23)) | (~ spl68_352 | spl68_471)), inference(forward_demodulation, [], [f3476, f2721])).
fof(f3476, plain, (~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | spl68_471), inference(avatar_component_clause, [], [f3474])).
fof(f3474, plain, (spl68_471 <=> (h2(op1(e12, e10)) = op2(h2(e12), h2(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl68_471])])).
fof(f8281, plain, (~ spl68_55 | ~ spl68_197 | ~ spl68_355 | spl68_486), inference(avatar_contradiction_clause, [], [f8280])).
fof(f8280, plain, ($false | (~ spl68_55 | ~ spl68_197 | ~ spl68_355 | spl68_486)), inference(subsumption_resolution, [], [f8279, f1764])).
fof(f8279, plain, (~ (e21 = op2(e22, e21)) | (~ spl68_55 | ~ spl68_355 | spl68_486)), inference(forward_demodulation, [], [f8278, f743])).
fof(f8278, plain, (~ (op2(e22, e21) = h2(e14)) | (~ spl68_55 | ~ spl68_355 | spl68_486)), inference(forward_demodulation, [], [f8277, f1117])).
fof(f1117, plain, ((e14 = op1(e12, e14)) | ~ spl68_55), inference(avatar_component_clause, [], [f1115])).
fof(f8277, plain, (~ (op2(e22, e21) = h2(op1(e12, e14))) | (~ spl68_355 | spl68_486)), inference(forward_demodulation, [], [f3536, f2736])).
fof(f3536, plain, (~ (h2(op1(e12, e14)) = op2(h2(e12), e21)) | spl68_486), inference(avatar_component_clause, [], [f3534])).
fof(f3534, plain, (spl68_486 <=> (h2(op1(e12, e14)) = op2(h2(e12), e21))), introduced(avatar_definition, [new_symbols(naming, [spl68_486])])).
fof(f8235, plain, (~ spl68_165 | ~ spl68_122 | ~ spl68_352 | spl68_463 | ~ spl68_490), inference(avatar_split_clause, [], [f8234, f3580, f3442, f2720, f1397, f1627])).
fof(f1397, plain, (spl68_122 <=> (op1(e10, e10) = e11)), introduced(avatar_definition, [new_symbols(naming, [spl68_122])])).
fof(f3442, plain, (spl68_463 <=> (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl68_463])])).
fof(f8234, plain, (~ (e24 = op2(e23, e23)) | (~ spl68_122 | ~ spl68_352 | spl68_463 | ~ spl68_490)), inference(forward_demodulation, [], [f8233, f3581])).
fof(f8233, plain, (~ (op2(e23, e23) = h2(e11)) | (~ spl68_122 | ~ spl68_352 | spl68_463)), inference(forward_demodulation, [], [f8232, f1399])).
fof(f1399, plain, ((op1(e10, e10) = e11) | ~ spl68_122), inference(avatar_component_clause, [], [f1397])).
fof(f8232, plain, (~ (op2(e23, e23) = h2(op1(e10, e10))) | (~ spl68_352 | spl68_463)), inference(forward_demodulation, [], [f3444, f2721])).
fof(f3444, plain, (~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10))) | spl68_463), inference(avatar_component_clause, [], [f3442])).
fof(f8185, plain, (~ spl68_201 | ~ spl68_202), inference(avatar_contradiction_clause, [], [f8184])).
fof(f8184, plain, ($false | (~ spl68_201 | ~ spl68_202)), inference(subsumption_resolution, [], [f8183, f545])).
fof(f545, plain, ~ (e20 = e21), inference(cnf_transformation, [], [f10])).
fof(f8183, plain, ((e20 = e21) | (~ spl68_201 | ~ spl68_202)), inference(forward_demodulation, [], [f1785, f1781])).
fof(f1785, plain, ((e21 = op2(e22, e20)) | ~ spl68_202), inference(avatar_component_clause, [], [f1783])).
fof(f1783, plain, (spl68_202 <=> (e21 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl68_202])])).
fof(f8182, plain, (~ spl68_237 | ~ spl68_240), inference(avatar_contradiction_clause, [], [f8181])).
fof(f8181, plain, ($false | (~ spl68_237 | ~ spl68_240)), inference(subsumption_resolution, [], [f8179, f551])).
fof(f8179, plain, ((e21 = e24) | (~ spl68_237 | ~ spl68_240)), inference(backward_demodulation, [], [f1944, f1932])).
fof(f8177, plain, (~ spl68_254 | spl68_366), inference(avatar_contradiction_clause, [], [f8176])).
fof(f8176, plain, ($false | (~ spl68_254 | spl68_366)), inference(subsumption_resolution, [], [f8174, f2793])).
fof(f2793, plain, (~ (e23 = h1(e11)) | spl68_366), inference(avatar_component_clause, [], [f2791])).
fof(f8174, plain, ((e23 = h1(e11)) | ~ spl68_254), inference(backward_demodulation, [], [f740, f2003])).
fof(f2003, plain, ((op2(e20, e20) = e23) | ~ spl68_254), inference(avatar_component_clause, [], [f2001])).
fof(f2001, plain, (spl68_254 <=> (op2(e20, e20) = e23)), introduced(avatar_definition, [new_symbols(naming, [spl68_254])])).
fof(f8167, plain, (~ spl68_376 | ~ spl68_122 | ~ spl68_314 | spl68_402), inference(avatar_split_clause, [], [f8166, f2983, f2524, f1397, f2841])).
fof(f8166, plain, (~ (e21 = h1(e11)) | (~ spl68_122 | ~ spl68_314 | spl68_402)), inference(forward_demodulation, [], [f8165, f2525])).
fof(f8165, plain, (~ (h1(e11) = h5(e11)) | (~ spl68_122 | spl68_402)), inference(forward_demodulation, [], [f2985, f1399])).
fof(f8144, plain, (~ spl68_76 | spl68_395), inference(avatar_contradiction_clause, [], [f8143])).
fof(f8143, plain, ($false | (~ spl68_76 | spl68_395)), inference(subsumption_resolution, [], [f8142, f2492])).
fof(f8142, plain, (~ (e20 = h5(e10)) | (~ spl68_76 | spl68_395)), inference(forward_demodulation, [], [f2957, f1206])).
fof(f2957, plain, (~ (e20 = h5(op1(e11, e14))) | spl68_395), inference(avatar_component_clause, [], [f2955])).
fof(f2955, plain, (spl68_395 <=> (e20 = h5(op1(e11, e14)))), introduced(avatar_definition, [new_symbols(naming, [spl68_395])])).
fof(f8141, plain, (~ spl68_71 | ~ spl68_201 | ~ spl68_310 | spl68_394), inference(avatar_contradiction_clause, [], [f8140])).
fof(f8140, plain, ($false | (~ spl68_71 | ~ spl68_201 | ~ spl68_310 | spl68_394)), inference(subsumption_resolution, [], [f8139, f1781])).
fof(f8139, plain, (~ (e20 = op2(e22, e20)) | (~ spl68_71 | ~ spl68_310 | spl68_394)), inference(forward_demodulation, [], [f8138, f2492])).
fof(f8138, plain, (~ (op2(e22, e20) = h5(e10)) | (~ spl68_71 | ~ spl68_310 | spl68_394)), inference(forward_demodulation, [], [f8137, f1185])).
fof(f8137, plain, (~ (op2(e22, e20) = h5(op1(e12, e10))) | (~ spl68_310 | spl68_394)), inference(forward_demodulation, [], [f2953, f2503])).
fof(f8131, plain, (~ spl68_83 | ~ spl68_213 | ~ spl68_310 | ~ spl68_314 | spl68_396), inference(avatar_contradiction_clause, [], [f8130])).
fof(f8130, plain, ($false | (~ spl68_83 | ~ spl68_213 | ~ spl68_310 | ~ spl68_314 | spl68_396)), inference(subsumption_resolution, [], [f8129, f1831])).
fof(f8129, plain, (~ (e22 = op2(e21, e23)) | (~ spl68_83 | ~ spl68_310 | ~ spl68_314 | spl68_396)), inference(forward_demodulation, [], [f8128, f2503])).
fof(f8128, plain, (~ (op2(e21, e23) = h5(e12)) | (~ spl68_83 | ~ spl68_314 | spl68_396)), inference(forward_demodulation, [], [f8127, f1235])).
fof(f8127, plain, (~ (op2(e21, e23) = h5(op1(e11, e13))) | (~ spl68_314 | spl68_396)), inference(forward_demodulation, [], [f2961, f2525])).
fof(f8126, plain, (~ spl68_23 | ~ spl68_153 | ~ spl68_310 | spl68_386), inference(avatar_contradiction_clause, [], [f8125])).
fof(f8125, plain, ($false | (~ spl68_23 | ~ spl68_153 | ~ spl68_310 | spl68_386)), inference(subsumption_resolution, [], [f8124, f1579])).
fof(f8124, plain, (~ (e22 = op2(e24, e20)) | (~ spl68_23 | ~ spl68_310 | spl68_386)), inference(forward_demodulation, [], [f8123, f2503])).
fof(f8123, plain, (~ (op2(e24, e20) = h5(e12)) | (~ spl68_23 | spl68_386)), inference(forward_demodulation, [], [f2921, f983])).
fof(f8116, plain, (~ spl68_99 | ~ spl68_229 | ~ spl68_314 | spl68_397), inference(avatar_contradiction_clause, [], [f8115])).
fof(f8115, plain, ($false | (~ spl68_99 | ~ spl68_229 | ~ spl68_314 | spl68_397)), inference(subsumption_resolution, [], [f8114, f1898])).
fof(f8114, plain, (~ (e23 = op2(e21, e20)) | (~ spl68_99 | ~ spl68_314 | spl68_397)), inference(forward_demodulation, [], [f8113, f2472])).
fof(f8113, plain, (~ (op2(e21, e20) = h5(e13)) | (~ spl68_99 | ~ spl68_314 | spl68_397)), inference(forward_demodulation, [], [f8112, f1302])).
fof(f8112, plain, (~ (op2(e21, e20) = h5(op1(e11, e10))) | (~ spl68_314 | spl68_397)), inference(forward_demodulation, [], [f2965, f2525])).
fof(f8111, plain, (~ spl68_15 | ~ spl68_145 | ~ spl68_310 | spl68_384), inference(avatar_contradiction_clause, [], [f8110])).
fof(f8110, plain, ($false | (~ spl68_15 | ~ spl68_145 | ~ spl68_310 | spl68_384)), inference(subsumption_resolution, [], [f8109, f1545])).
fof(f8109, plain, (~ (e24 = op2(e24, e22)) | (~ spl68_15 | ~ spl68_310 | spl68_384)), inference(forward_demodulation, [], [f8108, f758])).
fof(f8108, plain, (~ (op2(e24, e22) = h5(e14)) | (~ spl68_15 | ~ spl68_310 | spl68_384)), inference(forward_demodulation, [], [f8107, f949])).
fof(f8107, plain, (~ (op2(e24, e22) = h5(op1(e14, e12))) | (~ spl68_310 | spl68_384)), inference(forward_demodulation, [], [f2913, f2503])).
fof(f8102, plain, (~ spl68_104 | ~ spl68_234 | spl68_398), inference(avatar_contradiction_clause, [], [f8101])).
fof(f8101, plain, ($false | (~ spl68_104 | ~ spl68_234 | spl68_398)), inference(subsumption_resolution, [], [f8100, f1919])).
fof(f8100, plain, (~ (e23 = op2(e20, e24)) | (~ spl68_104 | spl68_398)), inference(forward_demodulation, [], [f8099, f2472])).
fof(f8099, plain, (~ (op2(e20, e24) = h5(e13)) | (~ spl68_104 | spl68_398)), inference(forward_demodulation, [], [f2969, f1323])).
fof(f1323, plain, ((e13 = op1(e10, e14)) | ~ spl68_104), inference(avatar_component_clause, [], [f1321])).
fof(f1321, plain, (spl68_104 <=> (e13 = op1(e10, e14))), introduced(avatar_definition, [new_symbols(naming, [spl68_104])])).
fof(f8098, plain, (~ spl68_6 | ~ spl68_136 | spl68_383), inference(avatar_contradiction_clause, [], [f8097])).
fof(f8097, plain, ($false | (~ spl68_6 | ~ spl68_136 | spl68_383)), inference(subsumption_resolution, [], [f8096, f1508])).
fof(f8096, plain, (~ (e20 = op2(e24, e23)) | (~ spl68_6 | spl68_383)), inference(forward_demodulation, [], [f8095, f2492])).
fof(f8095, plain, (~ (op2(e24, e23) = h5(e10)) | (~ spl68_6 | spl68_383)), inference(forward_demodulation, [], [f2909, f912])).
fof(f8094, plain, (~ spl68_2 | spl68_382), inference(avatar_contradiction_clause, [], [f8093])).
fof(f8093, plain, ($false | (~ spl68_2 | spl68_382)), inference(trivial_inequality_removal, [], [f8092])).
fof(f8092, plain, (~ (h5(e11) = h5(e11)) | (~ spl68_2 | spl68_382)), inference(forward_demodulation, [], [f2905, f895])).
fof(f2905, plain, (~ (h5(e11) = h5(op1(e14, e14))) | spl68_382), inference(avatar_component_clause, [], [f2903])).
fof(f2903, plain, (spl68_382 <=> (h5(e11) = h5(op1(e14, e14)))), introduced(avatar_definition, [new_symbols(naming, [spl68_382])])).
fof(f8085, plain, (~ spl68_111 | ~ spl68_241 | ~ spl68_310 | spl68_400), inference(avatar_contradiction_clause, [], [f8084])).
fof(f8084, plain, ($false | (~ spl68_111 | ~ spl68_241 | ~ spl68_310 | spl68_400)), inference(subsumption_resolution, [], [f8083, f1949])).
fof(f8083, plain, (~ (e20 = op2(e20, e22)) | (~ spl68_111 | ~ spl68_310 | spl68_400)), inference(forward_demodulation, [], [f8082, f2492])).
fof(f8082, plain, (~ (op2(e20, e22) = h5(e10)) | (~ spl68_111 | ~ spl68_310 | spl68_400)), inference(forward_demodulation, [], [f8081, f1353])).
fof(f8081, plain, (~ (op2(e20, e22) = h5(op1(e10, e12))) | (~ spl68_310 | spl68_400)), inference(forward_demodulation, [], [f2977, f2503])).
fof(f8080, plain, (~ spl68_63 | ~ spl68_193 | ~ spl68_310 | spl68_381), inference(avatar_contradiction_clause, [], [f8079])).
fof(f8079, plain, ($false | (~ spl68_63 | ~ spl68_193 | ~ spl68_310 | spl68_381)), inference(subsumption_resolution, [], [f8078, f1747])).
fof(f8078, plain, (~ (e22 = op2(e22, e22)) | (~ spl68_63 | ~ spl68_310 | spl68_381)), inference(forward_demodulation, [], [f8077, f2503])).
fof(f8077, plain, (~ (op2(e22, e22) = h5(e12)) | (~ spl68_63 | ~ spl68_310 | spl68_381)), inference(forward_demodulation, [], [f8076, f1151])).
fof(f8076, plain, (~ (op2(e22, e22) = h5(op1(e12, e12))) | (~ spl68_310 | spl68_381)), inference(forward_demodulation, [], [f2901, f2503])).
fof(f8065, plain, (~ spl68_67 | ~ spl68_197 | ~ spl68_310 | ~ spl68_314 | spl68_380), inference(avatar_contradiction_clause, [], [f8064])).
fof(f8064, plain, ($false | (~ spl68_67 | ~ spl68_197 | ~ spl68_310 | ~ spl68_314 | spl68_380)), inference(subsumption_resolution, [], [f8063, f1764])).
fof(f8063, plain, (~ (e21 = op2(e22, e21)) | (~ spl68_67 | ~ spl68_310 | ~ spl68_314 | spl68_380)), inference(forward_demodulation, [], [f8062, f2525])).
fof(f8062, plain, (~ (op2(e22, e21) = h5(e11)) | (~ spl68_67 | ~ spl68_310 | ~ spl68_314 | spl68_380)), inference(forward_demodulation, [], [f8061, f1168])).
fof(f8061, plain, (~ (op2(e22, e21) = h5(op1(e12, e11))) | (~ spl68_310 | ~ spl68_314 | spl68_380)), inference(forward_demodulation, [], [f8060, f2503])).
fof(f8060, plain, (~ (h5(op1(e12, e11)) = op2(h5(e12), e21)) | (~ spl68_314 | spl68_380)), inference(forward_demodulation, [], [f2897, f2525])).
fof(f8053, plain, (~ spl68_59 | ~ spl68_189 | ~ spl68_310 | spl68_393), inference(avatar_contradiction_clause, [], [f8052])).
fof(f8052, plain, ($false | (~ spl68_59 | ~ spl68_189 | ~ spl68_310 | spl68_393)), inference(subsumption_resolution, [], [f8051, f1730])).
fof(f8051, plain, (~ (e23 = op2(e22, e23)) | (~ spl68_59 | ~ spl68_310 | spl68_393)), inference(forward_demodulation, [], [f8050, f2472])).
fof(f8050, plain, (~ (op2(e22, e23) = h5(e13)) | (~ spl68_59 | ~ spl68_310 | spl68_393)), inference(forward_demodulation, [], [f8049, f1134])).
fof(f8049, plain, (~ (op2(e22, e23) = h5(op1(e12, e13))) | (~ spl68_310 | spl68_393)), inference(forward_demodulation, [], [f2949, f2503])).
fof(f8028, plain, (~ spl68_41 | ~ spl68_171 | ~ spl68_314 | spl68_390), inference(avatar_contradiction_clause, [], [f8027])).
fof(f8027, plain, ($false | (~ spl68_41 | ~ spl68_171 | ~ spl68_314 | spl68_390)), inference(subsumption_resolution, [], [f8026, f1655])).
fof(f8026, plain, (~ (e20 = op2(e23, e21)) | (~ spl68_41 | ~ spl68_314 | spl68_390)), inference(forward_demodulation, [], [f8025, f2492])).
fof(f8025, plain, (~ (op2(e23, e21) = h5(e10)) | (~ spl68_41 | ~ spl68_314 | spl68_390)), inference(forward_demodulation, [], [f8024, f1059])).
fof(f8024, plain, (~ (op2(e23, e21) = h5(op1(e13, e11))) | (~ spl68_314 | spl68_390)), inference(forward_demodulation, [], [f2937, f2525])).
fof(f2937, plain, (~ (h5(op1(e13, e11)) = op2(e23, h5(e11))) | spl68_390), inference(avatar_component_clause, [], [f2935])).
fof(f2935, plain, (spl68_390 <=> (h5(op1(e13, e11)) = op2(e23, h5(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl68_390])])).
fof(f7998, plain, (~ spl68_87 | ~ spl68_217 | ~ spl68_310 | ~ spl68_314 | spl68_379), inference(avatar_contradiction_clause, [], [f7997])).
fof(f7997, plain, ($false | (~ spl68_87 | ~ spl68_217 | ~ spl68_310 | ~ spl68_314 | spl68_379)), inference(subsumption_resolution, [], [f7996, f1848])).
fof(f7996, plain, (~ (e21 = op2(e21, e22)) | (~ spl68_87 | ~ spl68_310 | ~ spl68_314 | spl68_379)), inference(forward_demodulation, [], [f7995, f2525])).
fof(f7995, plain, (~ (op2(e21, e22) = h5(e11)) | (~ spl68_87 | ~ spl68_310 | ~ spl68_314 | spl68_379)), inference(forward_demodulation, [], [f7994, f1252])).
fof(f7994, plain, (~ (op2(e21, e22) = h5(op1(e11, e12))) | (~ spl68_310 | ~ spl68_314 | spl68_379)), inference(forward_demodulation, [], [f7993, f2525])).
fof(f7993, plain, (~ (h5(op1(e11, e12)) = op2(h5(e11), e22)) | (~ spl68_310 | spl68_379)), inference(forward_demodulation, [], [f2893, f2503])).
fof(f7964, plain, (~ spl68_50 | ~ spl68_180 | spl68_391), inference(avatar_contradiction_clause, [], [f7963])).
fof(f7963, plain, ($false | (~ spl68_50 | ~ spl68_180 | spl68_391)), inference(subsumption_resolution, [], [f7962, f1692])).
fof(f7962, plain, (~ (e24 = op2(e23, e20)) | (~ spl68_50 | spl68_391)), inference(forward_demodulation, [], [f7961, f758])).
fof(f7961, plain, (~ (op2(e23, e20) = h5(e14)) | (~ spl68_50 | spl68_391)), inference(forward_demodulation, [], [f2941, f1096])).
fof(f1096, plain, ((e14 = op1(e13, e10)) | ~ spl68_50), inference(avatar_component_clause, [], [f1094])).
fof(f1094, plain, (spl68_50 <=> (e14 = op1(e13, e10))), introduced(avatar_definition, [new_symbols(naming, [spl68_50])])).
fof(f7941, plain, (~ spl68_55 | ~ spl68_185 | ~ spl68_310 | spl68_392), inference(avatar_contradiction_clause, [], [f7940])).
fof(f7940, plain, ($false | (~ spl68_55 | ~ spl68_185 | ~ spl68_310 | spl68_392)), inference(subsumption_resolution, [], [f7939, f1713])).
fof(f7939, plain, (~ (e24 = op2(e22, e24)) | (~ spl68_55 | ~ spl68_310 | spl68_392)), inference(forward_demodulation, [], [f7938, f758])).
fof(f7938, plain, (~ (op2(e22, e24) = h5(e14)) | (~ spl68_55 | ~ spl68_310 | spl68_392)), inference(forward_demodulation, [], [f7937, f1117])).
fof(f7937, plain, (~ (op2(e22, e24) = h5(op1(e12, e14))) | (~ spl68_310 | spl68_392)), inference(forward_demodulation, [], [f2945, f2503])).
fof(f7901, plain, (~ spl68_248 | ~ spl68_118 | ~ spl68_310 | ~ spl68_314 | spl68_401), inference(avatar_split_clause, [], [f7900, f2979, f2524, f2502, f1380, f1976])).
fof(f7900, plain, (~ (e22 = op2(e20, e21)) | (~ spl68_118 | ~ spl68_310 | ~ spl68_314 | spl68_401)), inference(forward_demodulation, [], [f7899, f2503])).
fof(f7899, plain, (~ (op2(e20, e21) = h5(e12)) | (~ spl68_118 | ~ spl68_314 | spl68_401)), inference(forward_demodulation, [], [f7898, f1382])).
fof(f7898, plain, (~ (op2(e20, e21) = h5(op1(e10, e11))) | (~ spl68_314 | spl68_401)), inference(forward_demodulation, [], [f2981, f2525])).
fof(f7850, plain, (~ spl68_197 | ~ spl68_198), inference(avatar_contradiction_clause, [], [f7849])).
fof(f7849, plain, ($false | (~ spl68_197 | ~ spl68_198)), inference(subsumption_resolution, [], [f7848, f549])).
fof(f7848, plain, ((e21 = e22) | (~ spl68_197 | ~ spl68_198)), inference(forward_demodulation, [], [f1768, f1764])).
fof(f1768, plain, ((e22 = op2(e22, e21)) | ~ spl68_198), inference(avatar_component_clause, [], [f1766])).
fof(f1766, plain, (spl68_198 <=> (e22 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl68_198])])).
fof(f7831, plain, (~ spl68_28 | spl68_387), inference(avatar_contradiction_clause, [], [f7830])).
fof(f7830, plain, ($false | (~ spl68_28 | spl68_387)), inference(trivial_inequality_removal, [], [f7829])).
fof(f7829, plain, (~ (h5(e12) = h5(e12)) | (~ spl68_28 | spl68_387)), inference(forward_demodulation, [], [f2925, f1004])).
fof(f2925, plain, (~ (h5(e12) = h5(op1(e13, e14))) | spl68_387), inference(avatar_component_clause, [], [f2923])).
fof(f2923, plain, (spl68_387 <=> (h5(e12) = h5(op1(e13, e14)))), introduced(avatar_definition, [new_symbols(naming, [spl68_387])])).
fof(f7765, plain, (~ spl68_250 | ~ spl68_490), inference(avatar_split_clause, [], [f7147, f3580, f1984])).
fof(f1984, plain, (spl68_250 <=> (e24 = op2(e20, e21))), introduced(avatar_definition, [new_symbols(naming, [spl68_250])])).
fof(f7147, plain, (~ (e24 = op2(e20, e21)) | ~ spl68_490), inference(backward_demodulation, [], [f2440, f3581])).
fof(f2440, plain, ~ (op2(e20, e21) = h2(e11)), inference(backward_demodulation, [], [f445, f745])).
fof(f445, plain, ~ (op2(e20, e21) = op2(e21, e21)), inference(cnf_transformation, [], [f8])).
fof(f7758, plain, (~ spl68_162 | ~ spl68_165), inference(avatar_contradiction_clause, [], [f7757])).
fof(f7757, plain, ($false | (~ spl68_162 | ~ spl68_165)), inference(subsumption_resolution, [], [f7754, f551])).
fof(f7754, plain, ((e21 = e24) | (~ spl68_162 | ~ spl68_165)), inference(backward_demodulation, [], [f1629, f1617])).
fof(f1617, plain, ((e21 = op2(e23, e23)) | ~ spl68_162), inference(avatar_component_clause, [], [f1615])).
fof(f1615, plain, (spl68_162 <=> (e21 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl68_162])])).
fof(f7746, plain, (spl68_431 | ~ spl68_240 | ~ spl68_329), inference(avatar_split_clause, [], [f7745, f2604, f1942, f3205])).
fof(f7745, plain, ((e24 = h4(e12)) | (~ spl68_240 | ~ spl68_329)), inference(backward_demodulation, [], [f6174, f1944])).
fof(f6174, plain, ((op2(e20, e23) = h4(e12)) | ~ spl68_329), inference(backward_demodulation, [], [f2461, f2605])).
fof(f2461, plain, (h4(e12) = op2(h4(e13), e23)), inference(forward_demodulation, [], [f756, f757])).
fof(f757, plain, (op2(e23, op2(e23, e23)) = h4(e13)), inference(cnf_transformation, [], [f19])).
fof(f756, plain, (h4(e12) = op2(op2(e23, op2(e23, e23)), e23)), inference(cnf_transformation, [], [f19])).
fof(f7606, plain, (~ spl68_228 | ~ spl68_153), inference(avatar_split_clause, [], [f7605, f1577, f1892])).
fof(f1892, plain, (spl68_228 <=> (e22 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl68_228])])).
fof(f7605, plain, (~ (e22 = op2(e21, e20)) | ~ spl68_153), inference(forward_demodulation, [], [f441, f1579])).
fof(f441, plain, ~ (op2(e21, e20) = op2(e24, e20)), inference(cnf_transformation, [], [f8])).
fof(f7570, plain, (~ spl68_201 | ~ spl68_204), inference(avatar_contradiction_clause, [], [f7569])).
fof(f7569, plain, ($false | (~ spl68_201 | ~ spl68_204)), inference(subsumption_resolution, [], [f7568, f547])).
fof(f7568, plain, ((e20 = e23) | (~ spl68_201 | ~ spl68_204)), inference(forward_demodulation, [], [f1793, f1781])).
fof(f1793, plain, ((e23 = op2(e22, e20)) | ~ spl68_204), inference(avatar_component_clause, [], [f1791])).
fof(f1791, plain, (spl68_204 <=> (e23 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl68_204])])).
fof(f7555, plain, (~ spl68_255 | spl68_519), inference(avatar_contradiction_clause, [], [f7554])).
fof(f7554, plain, ($false | (~ spl68_255 | spl68_519)), inference(subsumption_resolution, [], [f7553, f3760])).
fof(f3760, plain, (~ (e24 = h1(e11)) | spl68_519), inference(avatar_component_clause, [], [f3758])).
fof(f7553, plain, ((e24 = h1(e11)) | ~ spl68_255), inference(backward_demodulation, [], [f740, f2007])).
fof(f2007, plain, ((op2(e20, e20) = e24) | ~ spl68_255), inference(avatar_component_clause, [], [f2005])).
fof(f2005, plain, (spl68_255 <=> (op2(e20, e20) = e24)), introduced(avatar_definition, [new_symbols(naming, [spl68_255])])).
fof(f7497, plain, (~ spl68_341 | ~ spl68_32 | ~ spl68_132 | spl68_449 | ~ spl68_450), inference(avatar_split_clause, [], [f7496, f3328, f3324, f1489, f1019, f2665])).
fof(f7496, plain, (~ (e21 = h3(e11)) | (~ spl68_32 | ~ spl68_132 | spl68_449 | ~ spl68_450)), inference(forward_demodulation, [], [f7495, f1491])).
fof(f7495, plain, (~ (op2(e24, e24) = h3(e11)) | (~ spl68_32 | spl68_449 | ~ spl68_450)), inference(forward_demodulation, [], [f7494, f1021])).
fof(f7494, plain, (~ (op2(e24, e24) = h3(op1(e13, e13))) | (spl68_449 | ~ spl68_450)), inference(forward_demodulation, [], [f3326, f3329])).
fof(f3326, plain, (~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | spl68_449), inference(avatar_component_clause, [], [f3324])).
fof(f7475, plain, (~ spl68_237 | ~ spl68_50 | ~ spl68_352 | ~ spl68_359 | spl68_475), inference(avatar_split_clause, [], [f7474, f3490, f2756, f2720, f1094, f1930])).
fof(f7474, plain, (~ (e21 = op2(e20, e23)) | (~ spl68_50 | ~ spl68_352 | ~ spl68_359 | spl68_475)), inference(forward_demodulation, [], [f7473, f743])).
fof(f7473, plain, (~ (op2(e20, e23) = h2(e14)) | (~ spl68_50 | ~ spl68_352 | ~ spl68_359 | spl68_475)), inference(forward_demodulation, [], [f7472, f1096])).
fof(f7472, plain, (~ (op2(e20, e23) = h2(op1(e13, e10))) | (~ spl68_352 | ~ spl68_359 | spl68_475)), inference(forward_demodulation, [], [f7471, f2757])).
fof(f7471, plain, (~ (h2(op1(e13, e10)) = op2(h2(e13), e23)) | (~ spl68_352 | spl68_475)), inference(forward_demodulation, [], [f3492, f2721])).
fof(f7446, plain, (spl68_355 | ~ spl68_248 | ~ spl68_359), inference(avatar_split_clause, [], [f7445, f2756, f1976, f2735])).
fof(f7445, plain, ((e22 = h2(e12)) | (~ spl68_248 | ~ spl68_359)), inference(forward_demodulation, [], [f7444, f1978])).
fof(f7444, plain, ((op2(e20, e21) = h2(e12)) | ~ spl68_359), inference(forward_demodulation, [], [f2439, f2757])).
fof(f2439, plain, (h2(e12) = op2(h2(e13), e21)), inference(forward_demodulation, [], [f746, f747])).
fof(f746, plain, (h2(e12) = op2(op2(e21, op2(e21, e21)), e21)), inference(cnf_transformation, [], [f17])).
fof(f7430, plain, (~ spl68_177 | ~ spl68_110 | ~ spl68_352 | ~ spl68_359 | spl68_466), inference(avatar_split_clause, [], [f7429, f3454, f2756, f2720, f1346, f1678])).
fof(f1346, plain, (spl68_110 <=> (e14 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl68_110])])).
fof(f7429, plain, (~ (e21 = op2(e23, e20)) | (~ spl68_110 | ~ spl68_352 | ~ spl68_359 | spl68_466)), inference(forward_demodulation, [], [f7428, f743])).
fof(f7428, plain, (~ (op2(e23, e20) = h2(e14)) | (~ spl68_110 | ~ spl68_352 | ~ spl68_359 | spl68_466)), inference(forward_demodulation, [], [f7427, f1348])).
fof(f1348, plain, ((e14 = op1(e10, e13)) | ~ spl68_110), inference(avatar_component_clause, [], [f1346])).
fof(f7427, plain, (~ (op2(e23, e20) = h2(op1(e10, e13))) | (~ spl68_352 | ~ spl68_359 | spl68_466)), inference(forward_demodulation, [], [f7426, f2721])).
fof(f7426, plain, (~ (h2(op1(e10, e13)) = op2(h2(e10), e20)) | (~ spl68_359 | spl68_466)), inference(forward_demodulation, [], [f3456, f2757])).
fof(f7425, plain, (~ spl68_255 | ~ spl68_32 | ~ spl68_359 | spl68_478 | ~ spl68_490), inference(avatar_split_clause, [], [f7424, f3580, f3502, f2756, f1019, f2005])).
fof(f7424, plain, (~ (op2(e20, e20) = e24) | (~ spl68_32 | ~ spl68_359 | spl68_478 | ~ spl68_490)), inference(forward_demodulation, [], [f7423, f3581])).
fof(f7423, plain, (~ (op2(e20, e20) = h2(e11)) | (~ spl68_32 | ~ spl68_359 | spl68_478)), inference(forward_demodulation, [], [f7422, f1021])).
fof(f7422, plain, (~ (op2(e20, e20) = h2(op1(e13, e13))) | (~ spl68_359 | spl68_478)), inference(forward_demodulation, [], [f3504, f2757])).
fof(f7360, plain, (~ spl68_241 | ~ spl68_245), inference(avatar_contradiction_clause, [], [f7359])).
fof(f7359, plain, ($false | (~ spl68_241 | ~ spl68_245)), inference(subsumption_resolution, [], [f7358, f548])).
fof(f548, plain, ~ (e20 = e24), inference(cnf_transformation, [], [f10])).
fof(f7358, plain, ((e20 = e24) | (~ spl68_241 | ~ spl68_245)), inference(forward_demodulation, [], [f1965, f1949])).
fof(f1965, plain, ((e24 = op2(e20, e22)) | ~ spl68_245), inference(avatar_component_clause, [], [f1963])).
fof(f1963, plain, (spl68_245 <=> (e24 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl68_245])])).
fof(f7353, plain, (~ spl68_341 | ~ spl68_217), inference(avatar_split_clause, [], [f7352, f1846, f2665])).
fof(f7352, plain, (~ (e21 = h3(e11)) | ~ spl68_217), inference(forward_demodulation, [], [f2452, f1848])).
fof(f2452, plain, ~ (op2(e21, e22) = h3(e11)), inference(backward_demodulation, [], [f459, f750])).
fof(f750, plain, (op2(e22, e22) = h3(e11)), inference(cnf_transformation, [], [f18])).
fof(f18, plain, ((op2(e22, op2(e22, e22)) = h3(e13)) & (h3(e12) = op2(op2(e22, op2(e22, e22)), e22)) & (op2(e22, e22) = h3(e11)) & (h3(e10) = op2(op2(e22, e22), e22)) & (e22 = h3(e14))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG052+1.p', ax18)).
fof(f459, plain, ~ (op2(e21, e22) = op2(e22, e22)), inference(cnf_transformation, [], [f8])).
fof(f7334, plain, (~ spl68_190 | ~ spl68_185), inference(avatar_split_clause, [], [f7333, f1711, f1732])).
fof(f1732, plain, (spl68_190 <=> (e24 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl68_190])])).
fof(f7333, plain, (~ (e24 = op2(e22, e23)) | ~ spl68_185), inference(forward_demodulation, [], [f514, f1713])).
fof(f514, plain, ~ (op2(e22, e23) = op2(e22, e24)), inference(cnf_transformation, [], [f8])).
fof(f7326, plain, (~ spl68_170 | ~ spl68_145), inference(avatar_split_clause, [], [f7325, f1543, f1648])).
fof(f7325, plain, (~ (e24 = op2(e23, e22)) | ~ spl68_145), inference(forward_demodulation, [], [f464, f1545])).
fof(f464, plain, ~ (op2(e23, e22) = op2(e24, e22)), inference(cnf_transformation, [], [f8])).
fof(f7324, plain, (~ spl68_169 | ~ spl68_39 | ~ spl68_310 | spl68_389), inference(avatar_split_clause, [], [f6326, f2931, f2502, f1048, f1644])).
fof(f6326, plain, (~ (e23 = op2(e23, e22)) | (~ spl68_39 | ~ spl68_310 | spl68_389)), inference(forward_demodulation, [], [f6325, f2472])).
fof(f6325, plain, (~ (op2(e23, e22) = h5(e13)) | (~ spl68_39 | ~ spl68_310 | spl68_389)), inference(forward_demodulation, [], [f6324, f1050])).
fof(f6324, plain, (~ (op2(e23, e22) = h5(op1(e13, e12))) | (~ spl68_310 | spl68_389)), inference(forward_demodulation, [], [f2933, f2503])).
fof(f7322, plain, (~ spl68_155 | ~ spl68_145), inference(avatar_split_clause, [], [f7321, f1543, f1585])).
fof(f7321, plain, (~ (e24 = op2(e24, e20)) | ~ spl68_145), inference(forward_demodulation, [], [f526, f1545])).
fof(f526, plain, ~ (op2(e24, e20) = op2(e24, e22)), inference(cnf_transformation, [], [f8])).
fof(f7217, plain, (~ spl68_252 | spl68_376), inference(avatar_contradiction_clause, [], [f7216])).
fof(f7216, plain, ($false | (~ spl68_252 | spl68_376)), inference(subsumption_resolution, [], [f7214, f2843])).
fof(f2843, plain, (~ (e21 = h1(e11)) | spl68_376), inference(avatar_component_clause, [], [f2841])).
fof(f7214, plain, ((e21 = h1(e11)) | ~ spl68_252), inference(backward_demodulation, [], [f740, f1995])).
fof(f7107, plain, (~ spl68_366 | ~ spl68_179), inference(avatar_split_clause, [], [f7106, f1686, f2791])).
fof(f7106, plain, (~ (e23 = h1(e11)) | ~ spl68_179), inference(forward_demodulation, [], [f2431, f1688])).
fof(f7097, plain, (~ spl68_248 | ~ spl68_243), inference(avatar_split_clause, [], [f6128, f1955, f1976])).
fof(f6128, plain, (~ (e22 = op2(e20, e21)) | ~ spl68_243), inference(forward_demodulation, [], [f489, f1957])).
fof(f489, plain, ~ (op2(e20, e21) = op2(e20, e22)), inference(cnf_transformation, [], [f8])).
fof(f7085, plain, (~ spl68_240 | ~ spl68_110 | spl68_399), inference(avatar_split_clause, [], [f7084, f2971, f1346, f1942])).
fof(f7084, plain, (~ (e24 = op2(e20, e23)) | (~ spl68_110 | spl68_399)), inference(forward_demodulation, [], [f7083, f758])).
fof(f7083, plain, (~ (op2(e20, e23) = h5(e14)) | (~ spl68_110 | spl68_399)), inference(forward_demodulation, [], [f2973, f1348])).
fof(f7057, plain, (~ spl68_152 | ~ spl68_314), inference(avatar_split_clause, [], [f4955, f2524, f1573])).
fof(f1573, plain, (spl68_152 <=> (e21 = op2(e24, e20))), introduced(avatar_definition, [new_symbols(naming, [spl68_152])])).
fof(f4955, plain, (~ (e21 = op2(e24, e20)) | ~ spl68_314), inference(backward_demodulation, [], [f2484, f2525])).
fof(f2484, plain, ~ (op2(e24, e20) = h5(e11)), inference(backward_demodulation, [], [f528, f760])).
fof(f528, plain, ~ (op2(e24, e20) = op2(e24, e24)), inference(cnf_transformation, [], [f8])).
fof(f7053, plain, (~ spl68_143 | ~ spl68_243), inference(avatar_split_clause, [], [f7052, f1955, f1535])).
fof(f1535, plain, (spl68_143 <=> (e22 = op2(e24, e22))), introduced(avatar_definition, [new_symbols(naming, [spl68_143])])).
fof(f7052, plain, (~ (e22 = op2(e24, e22)) | ~ spl68_243), inference(forward_demodulation, [], [f458, f1957])).
fof(f7046, plain, (~ spl68_41 | ~ spl68_45), inference(avatar_contradiction_clause, [], [f7045])).
fof(f7045, plain, ($false | (~ spl68_41 | ~ spl68_45)), inference(subsumption_resolution, [], [f7044, f538])).
fof(f538, plain, ~ (e10 = e14), inference(cnf_transformation, [], [f9])).
fof(f7044, plain, ((e10 = e14) | (~ spl68_41 | ~ spl68_45)), inference(forward_demodulation, [], [f1075, f1059])).
fof(f1075, plain, ((e14 = op1(e13, e11)) | ~ spl68_45), inference(avatar_component_clause, [], [f1073])).
fof(f1073, plain, (spl68_45 <=> (e14 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl68_45])])).
fof(f7014, plain, (~ spl68_94 | ~ spl68_19), inference(avatar_split_clause, [], [f6984, f964, f1279])).
fof(f1279, plain, (spl68_94 <=> (e13 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl68_94])])).
fof(f6984, plain, (~ (e13 = op1(e11, e11)) | ~ spl68_19), inference(forward_demodulation, [], [f351, f966])).
fof(f351, plain, ~ (op1(e11, e11) = op1(e14, e11)), inference(cnf_transformation, [], [f7])).
fof(f6980, plain, (~ spl68_91 | ~ spl68_76), inference(avatar_split_clause, [], [f6979, f1204, f1267])).
fof(f1267, plain, (spl68_91 <=> (e10 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl68_91])])).
fof(f6979, plain, (~ (e10 = op1(e11, e11)) | ~ spl68_76), inference(forward_demodulation, [], [f401, f1206])).
fof(f401, plain, ~ (op1(e11, e11) = op1(e11, e14)), inference(cnf_transformation, [], [f7])).
fof(f6974, plain, (~ spl68_104 | ~ spl68_170 | ~ spl68_337 | ~ spl68_450 | spl68_459), inference(avatar_contradiction_clause, [], [f6973])).
fof(f6973, plain, ($false | (~ spl68_104 | ~ spl68_170 | ~ spl68_337 | ~ spl68_450 | spl68_459)), inference(subsumption_resolution, [], [f6972, f1650])).
fof(f6972, plain, (~ (e24 = op2(e23, e22)) | (~ spl68_104 | ~ spl68_337 | ~ spl68_450 | spl68_459)), inference(forward_demodulation, [], [f6971, f3329])).
fof(f6971, plain, (~ (op2(e23, e22) = h3(e13)) | (~ spl68_104 | ~ spl68_337 | spl68_459)), inference(forward_demodulation, [], [f6970, f1323])).
fof(f6970, plain, (~ (op2(e23, e22) = h3(op1(e10, e14))) | (~ spl68_337 | spl68_459)), inference(forward_demodulation, [], [f3366, f2645])).
fof(f3366, plain, (~ (h3(op1(e10, e14)) = op2(h3(e10), e22)) | spl68_459), inference(avatar_component_clause, [], [f3364])).
fof(f3364, plain, (spl68_459 <=> (h3(op1(e10, e14)) = op2(h3(e10), e22))), introduced(avatar_definition, [new_symbols(naming, [spl68_459])])).
fof(f6966, plain, (~ spl68_76 | spl68_458), inference(avatar_contradiction_clause, [], [f6965])).
fof(f6965, plain, ($false | (~ spl68_76 | spl68_458)), inference(trivial_inequality_removal, [], [f6964])).
fof(f6964, plain, (~ (h3(e10) = h3(e10)) | (~ spl68_76 | spl68_458)), inference(forward_demodulation, [], [f3362, f1206])).
fof(f3362, plain, (~ (h3(e10) = h3(op1(e11, e14))) | spl68_458), inference(avatar_component_clause, [], [f3360])).
fof(f3360, plain, (spl68_458 <=> (h3(e10) = h3(op1(e11, e14)))), introduced(avatar_definition, [new_symbols(naming, [spl68_458])])).
fof(f6963, plain, (~ spl68_55 | ~ spl68_243 | ~ spl68_345 | spl68_457), inference(avatar_contradiction_clause, [], [f6962])).
fof(f6962, plain, ($false | (~ spl68_55 | ~ spl68_243 | ~ spl68_345 | spl68_457)), inference(subsumption_resolution, [], [f6961, f1957])).
fof(f6961, plain, (~ (e22 = op2(e20, e22)) | (~ spl68_55 | ~ spl68_345 | spl68_457)), inference(forward_demodulation, [], [f6960, f748])).
fof(f748, plain, (e22 = h3(e14)), inference(cnf_transformation, [], [f18])).
fof(f6960, plain, (~ (op2(e20, e22) = h3(e14)) | (~ spl68_55 | ~ spl68_345 | spl68_457)), inference(forward_demodulation, [], [f6959, f1117])).
fof(f6959, plain, (~ (op2(e20, e22) = h3(op1(e12, e14))) | (~ spl68_345 | spl68_457)), inference(forward_demodulation, [], [f3358, f2686])).
fof(f2686, plain, ((e20 = h3(e12)) | ~ spl68_345), inference(avatar_component_clause, [], [f2685])).
fof(f2685, plain, (spl68_345 <=> (e20 = h3(e12))), introduced(avatar_definition, [new_symbols(naming, [spl68_345])])).
fof(f3358, plain, (~ (h3(op1(e12, e14)) = op2(h3(e12), e22)) | spl68_457), inference(avatar_component_clause, [], [f3356])).
fof(f3356, plain, (spl68_457 <=> (h3(op1(e12, e14)) = op2(h3(e12), e22))), introduced(avatar_definition, [new_symbols(naming, [spl68_457])])).
fof(f6955, plain, (~ spl68_28 | spl68_456), inference(avatar_contradiction_clause, [], [f6954])).
fof(f6954, plain, ($false | (~ spl68_28 | spl68_456)), inference(trivial_inequality_removal, [], [f6953])).
fof(f6953, plain, (~ (h3(e12) = h3(e12)) | (~ spl68_28 | spl68_456)), inference(forward_demodulation, [], [f3354, f1004])).
fof(f3354, plain, (~ (h3(e12) = h3(op1(e13, e14))) | spl68_456), inference(avatar_component_clause, [], [f3352])).
fof(f3352, plain, (spl68_456 <=> (h3(e12) = h3(op1(e13, e14)))), introduced(avatar_definition, [new_symbols(naming, [spl68_456])])).
fof(f6952, plain, (~ spl68_23 | ~ spl68_186 | ~ spl68_337 | ~ spl68_345 | spl68_455), inference(avatar_contradiction_clause, [], [f6951])).
fof(f6951, plain, ($false | (~ spl68_23 | ~ spl68_186 | ~ spl68_337 | ~ spl68_345 | spl68_455)), inference(subsumption_resolution, [], [f6950, f1718])).
fof(f6950, plain, (~ (e20 = op2(e22, e23)) | (~ spl68_23 | ~ spl68_337 | ~ spl68_345 | spl68_455)), inference(forward_demodulation, [], [f6949, f2686])).
fof(f6949, plain, (~ (op2(e22, e23) = h3(e12)) | (~ spl68_23 | ~ spl68_337 | spl68_455)), inference(forward_demodulation, [], [f6948, f983])).
fof(f6948, plain, (~ (op2(e22, e23) = h3(op1(e14, e10))) | (~ spl68_337 | spl68_455)), inference(forward_demodulation, [], [f3350, f2645])).
fof(f3350, plain, (~ (h3(op1(e14, e10)) = op2(e22, h3(e10))) | spl68_455), inference(avatar_component_clause, [], [f3348])).
fof(f3348, plain, (spl68_455 <=> (h3(op1(e14, e10)) = op2(e22, h3(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl68_455])])).
fof(f6944, plain, (~ spl68_19 | spl68_454), inference(avatar_contradiction_clause, [], [f6943])).
fof(f6943, plain, ($false | (~ spl68_19 | spl68_454)), inference(trivial_inequality_removal, [], [f6942])).
fof(f6942, plain, (~ (h3(e13) = h3(e13)) | (~ spl68_19 | spl68_454)), inference(forward_demodulation, [], [f3346, f966])).
fof(f3346, plain, (~ (h3(e13) = h3(op1(e14, e11))) | spl68_454), inference(avatar_component_clause, [], [f3344])).
fof(f3344, plain, (spl68_454 <=> (h3(e13) = h3(op1(e14, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl68_454])])).
fof(f6938, plain, (~ spl68_15 | ~ spl68_203 | ~ spl68_345 | spl68_453), inference(avatar_contradiction_clause, [], [f6937])).
fof(f6937, plain, ($false | (~ spl68_15 | ~ spl68_203 | ~ spl68_345 | spl68_453)), inference(subsumption_resolution, [], [f6936, f1789])).
fof(f6936, plain, (~ (e22 = op2(e22, e20)) | (~ spl68_15 | ~ spl68_345 | spl68_453)), inference(forward_demodulation, [], [f6935, f748])).
fof(f6935, plain, (~ (op2(e22, e20) = h3(e14)) | (~ spl68_15 | ~ spl68_345 | spl68_453)), inference(forward_demodulation, [], [f6934, f949])).
fof(f6934, plain, (~ (op2(e22, e20) = h3(op1(e14, e12))) | (~ spl68_345 | spl68_453)), inference(forward_demodulation, [], [f3342, f2686])).
fof(f3342, plain, (~ (h3(op1(e14, e12)) = op2(e22, h3(e12))) | spl68_453), inference(avatar_component_clause, [], [f3340])).
fof(f3340, plain, (spl68_453 <=> (h3(op1(e14, e12)) = op2(e22, h3(e12)))), introduced(avatar_definition, [new_symbols(naming, [spl68_453])])).
fof(f6933, plain, (~ spl68_6 | ~ spl68_184 | ~ spl68_337 | ~ spl68_450 | spl68_452), inference(avatar_contradiction_clause, [], [f6932])).
fof(f6932, plain, ($false | (~ spl68_6 | ~ spl68_184 | ~ spl68_337 | ~ spl68_450 | spl68_452)), inference(subsumption_resolution, [], [f6931, f1709])).
fof(f6931, plain, (~ (e23 = op2(e22, e24)) | (~ spl68_6 | ~ spl68_337 | ~ spl68_450 | spl68_452)), inference(forward_demodulation, [], [f6930, f2645])).
fof(f6930, plain, (~ (op2(e22, e24) = h3(e10)) | (~ spl68_6 | ~ spl68_450 | spl68_452)), inference(forward_demodulation, [], [f6929, f912])).
fof(f6929, plain, (~ (op2(e22, e24) = h3(op1(e14, e13))) | (~ spl68_450 | spl68_452)), inference(forward_demodulation, [], [f3338, f3329])).
fof(f3338, plain, (~ (h3(op1(e14, e13)) = op2(e22, h3(e13))) | spl68_452), inference(avatar_component_clause, [], [f3336])).
fof(f3336, plain, (spl68_452 <=> (h3(op1(e14, e13)) = op2(e22, h3(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl68_452])])).
fof(f6925, plain, (~ spl68_2 | spl68_451), inference(avatar_contradiction_clause, [], [f6924])).
fof(f6924, plain, ($false | (~ spl68_2 | spl68_451)), inference(trivial_inequality_removal, [], [f6923])).
fof(f6923, plain, (~ (h3(e11) = h3(e11)) | (~ spl68_2 | spl68_451)), inference(forward_demodulation, [], [f3334, f895])).
fof(f3334, plain, (~ (h3(e11) = h3(op1(e14, e14))) | spl68_451), inference(avatar_component_clause, [], [f3332])).
fof(f3332, plain, (spl68_451 <=> (h3(e11) = h3(op1(e14, e14)))), introduced(avatar_definition, [new_symbols(naming, [spl68_451])])).
fof(f6909, plain, (~ spl68_39 | ~ spl68_155 | ~ spl68_345 | spl68_448 | ~ spl68_450), inference(avatar_contradiction_clause, [], [f6908])).
fof(f6908, plain, ($false | (~ spl68_39 | ~ spl68_155 | ~ spl68_345 | spl68_448 | ~ spl68_450)), inference(subsumption_resolution, [], [f6907, f1587])).
fof(f6907, plain, (~ (e24 = op2(e24, e20)) | (~ spl68_39 | ~ spl68_345 | spl68_448 | ~ spl68_450)), inference(forward_demodulation, [], [f6906, f3329])).
fof(f6906, plain, (~ (op2(e24, e20) = h3(e13)) | (~ spl68_39 | ~ spl68_345 | spl68_448 | ~ spl68_450)), inference(forward_demodulation, [], [f6905, f1050])).
fof(f6905, plain, (~ (op2(e24, e20) = h3(op1(e13, e12))) | (~ spl68_345 | spl68_448 | ~ spl68_450)), inference(forward_demodulation, [], [f6904, f3329])).
fof(f6904, plain, (~ (h3(op1(e13, e12)) = op2(h3(e13), e20)) | (~ spl68_345 | spl68_448)), inference(forward_demodulation, [], [f3322, f2686])).
fof(f3322, plain, (~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | spl68_448), inference(avatar_component_clause, [], [f3320])).
fof(f3320, plain, (spl68_448 <=> (h3(op1(e13, e12)) = op2(h3(e13), h3(e12)))), introduced(avatar_definition, [new_symbols(naming, [spl68_448])])).
fof(f6899, plain, (~ spl68_41 | ~ spl68_149 | ~ spl68_337 | ~ spl68_341 | spl68_447 | ~ spl68_450), inference(avatar_contradiction_clause, [], [f6898])).
fof(f6898, plain, ($false | (~ spl68_41 | ~ spl68_149 | ~ spl68_337 | ~ spl68_341 | spl68_447 | ~ spl68_450)), inference(subsumption_resolution, [], [f6897, f1562])).
fof(f6897, plain, (~ (e23 = op2(e24, e21)) | (~ spl68_41 | ~ spl68_337 | ~ spl68_341 | spl68_447 | ~ spl68_450)), inference(forward_demodulation, [], [f6896, f2645])).
fof(f6896, plain, (~ (op2(e24, e21) = h3(e10)) | (~ spl68_41 | ~ spl68_341 | spl68_447 | ~ spl68_450)), inference(forward_demodulation, [], [f6895, f1059])).
fof(f6895, plain, (~ (op2(e24, e21) = h3(op1(e13, e11))) | (~ spl68_341 | spl68_447 | ~ spl68_450)), inference(forward_demodulation, [], [f6894, f3329])).
fof(f6894, plain, (~ (h3(op1(e13, e11)) = op2(h3(e13), e21)) | (~ spl68_341 | spl68_447)), inference(forward_demodulation, [], [f3318, f2666])).
fof(f3318, plain, (~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | spl68_447), inference(avatar_component_clause, [], [f3316])).
fof(f3316, plain, (spl68_447 <=> (h3(op1(e13, e11)) = op2(h3(e13), h3(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl68_447])])).
fof(f6889, plain, (~ spl68_50 | ~ spl68_138 | ~ spl68_337 | spl68_446 | ~ spl68_450), inference(avatar_contradiction_clause, [], [f6888])).
fof(f6888, plain, ($false | (~ spl68_50 | ~ spl68_138 | ~ spl68_337 | spl68_446 | ~ spl68_450)), inference(subsumption_resolution, [], [f6887, f1516])).
fof(f6887, plain, (~ (e22 = op2(e24, e23)) | (~ spl68_50 | ~ spl68_337 | spl68_446 | ~ spl68_450)), inference(forward_demodulation, [], [f6886, f748])).
fof(f6886, plain, (~ (op2(e24, e23) = h3(e14)) | (~ spl68_50 | ~ spl68_337 | spl68_446 | ~ spl68_450)), inference(forward_demodulation, [], [f6885, f1096])).
fof(f6885, plain, (~ (op2(e24, e23) = h3(op1(e13, e10))) | (~ spl68_337 | spl68_446 | ~ spl68_450)), inference(forward_demodulation, [], [f6884, f3329])).
fof(f6884, plain, (~ (h3(op1(e13, e10)) = op2(h3(e13), e23)) | (~ spl68_337 | spl68_446)), inference(forward_demodulation, [], [f3314, f2645])).
fof(f3314, plain, (~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | spl68_446), inference(avatar_component_clause, [], [f3312])).
fof(f3312, plain, (spl68_446 <=> (h3(op1(e13, e10)) = op2(h3(e13), h3(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl68_446])])).
fof(f6880, plain, (~ spl68_59 | ~ spl68_235 | ~ spl68_345 | spl68_445 | ~ spl68_450), inference(avatar_contradiction_clause, [], [f6879])).
fof(f6879, plain, ($false | (~ spl68_59 | ~ spl68_235 | ~ spl68_345 | spl68_445 | ~ spl68_450)), inference(subsumption_resolution, [], [f6878, f1923])).
fof(f6878, plain, (~ (e24 = op2(e20, e24)) | (~ spl68_59 | ~ spl68_345 | spl68_445 | ~ spl68_450)), inference(forward_demodulation, [], [f6877, f3329])).
fof(f6877, plain, (~ (op2(e20, e24) = h3(e13)) | (~ spl68_59 | ~ spl68_345 | spl68_445 | ~ spl68_450)), inference(forward_demodulation, [], [f6876, f1134])).
fof(f6876, plain, (~ (op2(e20, e24) = h3(op1(e12, e13))) | (~ spl68_345 | spl68_445 | ~ spl68_450)), inference(forward_demodulation, [], [f6875, f2686])).
fof(f6875, plain, (~ (h3(op1(e12, e13)) = op2(h3(e12), e24)) | (spl68_445 | ~ spl68_450)), inference(forward_demodulation, [], [f3310, f3329])).
fof(f3310, plain, (~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | spl68_445), inference(avatar_component_clause, [], [f3308])).
fof(f3308, plain, (spl68_445 <=> (h3(op1(e12, e13)) = op2(h3(e12), h3(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl68_445])])).
fof(f6870, plain, (~ spl68_63 | ~ spl68_251 | ~ spl68_345 | spl68_444), inference(avatar_contradiction_clause, [], [f6869])).
fof(f6869, plain, ($false | (~ spl68_63 | ~ spl68_251 | ~ spl68_345 | spl68_444)), inference(subsumption_resolution, [], [f6868, f1991])).
fof(f6868, plain, (~ (e20 = op2(e20, e20)) | (~ spl68_63 | ~ spl68_345 | spl68_444)), inference(forward_demodulation, [], [f6867, f2686])).
fof(f6867, plain, (~ (op2(e20, e20) = h3(e12)) | (~ spl68_63 | ~ spl68_345 | spl68_444)), inference(forward_demodulation, [], [f6866, f1151])).
fof(f6866, plain, (~ (op2(e20, e20) = h3(op1(e12, e12))) | (~ spl68_345 | spl68_444)), inference(forward_demodulation, [], [f3306, f2686])).
fof(f3306, plain, (~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | spl68_444), inference(avatar_component_clause, [], [f3304])).
fof(f3304, plain, (spl68_444 <=> (h3(op1(e12, e12)) = op2(h3(e12), h3(e12)))), introduced(avatar_definition, [new_symbols(naming, [spl68_444])])).
fof(f6861, plain, (~ spl68_67 | ~ spl68_247 | ~ spl68_341 | ~ spl68_345 | spl68_443), inference(avatar_contradiction_clause, [], [f6860])).
fof(f6860, plain, ($false | (~ spl68_67 | ~ spl68_247 | ~ spl68_341 | ~ spl68_345 | spl68_443)), inference(subsumption_resolution, [], [f6859, f1974])).
fof(f6859, plain, (~ (e21 = op2(e20, e21)) | (~ spl68_67 | ~ spl68_341 | ~ spl68_345 | spl68_443)), inference(forward_demodulation, [], [f6858, f2666])).
fof(f6858, plain, (~ (op2(e20, e21) = h3(e11)) | (~ spl68_67 | ~ spl68_341 | ~ spl68_345 | spl68_443)), inference(forward_demodulation, [], [f6857, f1168])).
fof(f6857, plain, (~ (op2(e20, e21) = h3(op1(e12, e11))) | (~ spl68_341 | ~ spl68_345 | spl68_443)), inference(forward_demodulation, [], [f6856, f2686])).
fof(f6856, plain, (~ (h3(op1(e12, e11)) = op2(h3(e12), e21)) | (~ spl68_341 | spl68_443)), inference(forward_demodulation, [], [f3302, f2666])).
fof(f3302, plain, (~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | spl68_443), inference(avatar_component_clause, [], [f3300])).
fof(f3300, plain, (spl68_443 <=> (h3(op1(e12, e11)) = op2(h3(e12), h3(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl68_443])])).
fof(f6851, plain, (~ spl68_71 | ~ spl68_239 | ~ spl68_337 | ~ spl68_345 | spl68_442), inference(avatar_contradiction_clause, [], [f6850])).
fof(f6850, plain, ($false | (~ spl68_71 | ~ spl68_239 | ~ spl68_337 | ~ spl68_345 | spl68_442)), inference(subsumption_resolution, [], [f6849, f1940])).
fof(f6849, plain, (~ (e23 = op2(e20, e23)) | (~ spl68_71 | ~ spl68_337 | ~ spl68_345 | spl68_442)), inference(forward_demodulation, [], [f6848, f2645])).
fof(f6848, plain, (~ (op2(e20, e23) = h3(e10)) | (~ spl68_71 | ~ spl68_337 | ~ spl68_345 | spl68_442)), inference(forward_demodulation, [], [f6847, f1185])).
fof(f6847, plain, (~ (op2(e20, e23) = h3(op1(e12, e10))) | (~ spl68_337 | ~ spl68_345 | spl68_442)), inference(forward_demodulation, [], [f6846, f2686])).
fof(f6846, plain, (~ (h3(op1(e12, e10)) = op2(h3(e12), e23)) | (~ spl68_337 | spl68_442)), inference(forward_demodulation, [], [f3298, f2645])).
fof(f3298, plain, (~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | spl68_442), inference(avatar_component_clause, [], [f3296])).
fof(f3296, plain, (spl68_442 <=> (h3(op1(e12, e10)) = op2(h3(e12), h3(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl68_442])])).
fof(f6841, plain, (~ spl68_83 | ~ spl68_206 | ~ spl68_341 | ~ spl68_345 | spl68_441 | ~ spl68_450), inference(avatar_contradiction_clause, [], [f6840])).
fof(f6840, plain, ($false | (~ spl68_83 | ~ spl68_206 | ~ spl68_341 | ~ spl68_345 | spl68_441 | ~ spl68_450)), inference(subsumption_resolution, [], [f6839, f1802])).
fof(f6839, plain, (~ (e20 = op2(e21, e24)) | (~ spl68_83 | ~ spl68_341 | ~ spl68_345 | spl68_441 | ~ spl68_450)), inference(forward_demodulation, [], [f6838, f2686])).
fof(f6838, plain, (~ (op2(e21, e24) = h3(e12)) | (~ spl68_83 | ~ spl68_341 | spl68_441 | ~ spl68_450)), inference(forward_demodulation, [], [f6837, f1235])).
fof(f6837, plain, (~ (op2(e21, e24) = h3(op1(e11, e13))) | (~ spl68_341 | spl68_441 | ~ spl68_450)), inference(forward_demodulation, [], [f6836, f2666])).
fof(f6836, plain, (~ (h3(op1(e11, e13)) = op2(h3(e11), e24)) | (spl68_441 | ~ spl68_450)), inference(forward_demodulation, [], [f3294, f3329])).
fof(f3294, plain, (~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | spl68_441), inference(avatar_component_clause, [], [f3292])).
fof(f3292, plain, (spl68_441 <=> (h3(op1(e11, e13)) = op2(h3(e11), h3(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl68_441])])).
fof(f6832, plain, (~ spl68_87 | ~ spl68_227 | ~ spl68_341 | ~ spl68_345 | spl68_440), inference(avatar_contradiction_clause, [], [f6831])).
fof(f6831, plain, ($false | (~ spl68_87 | ~ spl68_227 | ~ spl68_341 | ~ spl68_345 | spl68_440)), inference(subsumption_resolution, [], [f6830, f1890])).
fof(f6830, plain, (~ (e21 = op2(e21, e20)) | (~ spl68_87 | ~ spl68_341 | ~ spl68_345 | spl68_440)), inference(forward_demodulation, [], [f6829, f2666])).
fof(f6829, plain, (~ (op2(e21, e20) = h3(e11)) | (~ spl68_87 | ~ spl68_341 | ~ spl68_345 | spl68_440)), inference(forward_demodulation, [], [f6828, f1252])).
fof(f6828, plain, (~ (op2(e21, e20) = h3(op1(e11, e12))) | (~ spl68_341 | ~ spl68_345 | spl68_440)), inference(forward_demodulation, [], [f6827, f2666])).
fof(f6827, plain, (~ (h3(op1(e11, e12)) = op2(h3(e11), e20)) | (~ spl68_345 | spl68_440)), inference(forward_demodulation, [], [f3290, f2686])).
fof(f3290, plain, (~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | spl68_440), inference(avatar_component_clause, [], [f3288])).
fof(f3288, plain, (spl68_440 <=> (h3(op1(e11, e12)) = op2(h3(e11), h3(e12)))), introduced(avatar_definition, [new_symbols(naming, [spl68_440])])).
fof(f6822, plain, (~ spl68_95 | ~ spl68_223 | ~ spl68_341 | spl68_439), inference(avatar_contradiction_clause, [], [f6821])).
fof(f6821, plain, ($false | (~ spl68_95 | ~ spl68_223 | ~ spl68_341 | spl68_439)), inference(subsumption_resolution, [], [f6820, f1873])).
fof(f6820, plain, (~ (e22 = op2(e21, e21)) | (~ spl68_95 | ~ spl68_341 | spl68_439)), inference(forward_demodulation, [], [f6819, f748])).
fof(f6819, plain, (~ (op2(e21, e21) = h3(e14)) | (~ spl68_95 | ~ spl68_341 | spl68_439)), inference(forward_demodulation, [], [f6818, f1285])).
fof(f6818, plain, (~ (op2(e21, e21) = h3(op1(e11, e11))) | (~ spl68_341 | spl68_439)), inference(forward_demodulation, [], [f3286, f2666])).
fof(f3286, plain, (~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | spl68_439), inference(avatar_component_clause, [], [f3284])).
fof(f3284, plain, (spl68_439 <=> (h3(op1(e11, e11)) = op2(h3(e11), h3(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl68_439])])).
fof(f6813, plain, (~ spl68_99 | ~ spl68_215 | ~ spl68_337 | ~ spl68_341 | spl68_438 | ~ spl68_450), inference(avatar_contradiction_clause, [], [f6812])).
fof(f6812, plain, ($false | (~ spl68_99 | ~ spl68_215 | ~ spl68_337 | ~ spl68_341 | spl68_438 | ~ spl68_450)), inference(subsumption_resolution, [], [f6811, f1839])).
fof(f6811, plain, (~ (e24 = op2(e21, e23)) | (~ spl68_99 | ~ spl68_337 | ~ spl68_341 | spl68_438 | ~ spl68_450)), inference(forward_demodulation, [], [f6810, f3329])).
fof(f6810, plain, (~ (op2(e21, e23) = h3(e13)) | (~ spl68_99 | ~ spl68_337 | ~ spl68_341 | spl68_438)), inference(forward_demodulation, [], [f6809, f1302])).
fof(f6809, plain, (~ (op2(e21, e23) = h3(op1(e11, e10))) | (~ spl68_337 | ~ spl68_341 | spl68_438)), inference(forward_demodulation, [], [f6808, f2666])).
fof(f6808, plain, (~ (h3(op1(e11, e10)) = op2(h3(e11), e23)) | (~ spl68_337 | spl68_438)), inference(forward_demodulation, [], [f3282, f2645])).
fof(f3282, plain, (~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | spl68_438), inference(avatar_component_clause, [], [f3280])).
fof(f3280, plain, (spl68_438 <=> (h3(op1(e11, e10)) = op2(h3(e11), h3(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl68_438])])).
fof(f6803, plain, (~ spl68_111 | ~ spl68_179 | ~ spl68_337 | ~ spl68_345 | spl68_436), inference(avatar_contradiction_clause, [], [f6802])).
fof(f6802, plain, ($false | (~ spl68_111 | ~ spl68_179 | ~ spl68_337 | ~ spl68_345 | spl68_436)), inference(subsumption_resolution, [], [f6801, f1688])).
fof(f6801, plain, (~ (e23 = op2(e23, e20)) | (~ spl68_111 | ~ spl68_337 | ~ spl68_345 | spl68_436)), inference(forward_demodulation, [], [f6800, f2645])).
fof(f6800, plain, (~ (op2(e23, e20) = h3(e10)) | (~ spl68_111 | ~ spl68_337 | ~ spl68_345 | spl68_436)), inference(forward_demodulation, [], [f6799, f1353])).
fof(f6799, plain, (~ (op2(e23, e20) = h3(op1(e10, e12))) | (~ spl68_337 | ~ spl68_345 | spl68_436)), inference(forward_demodulation, [], [f6798, f2645])).
fof(f6798, plain, (~ (h3(op1(e10, e12)) = op2(h3(e10), e20)) | (~ spl68_345 | spl68_436)), inference(forward_demodulation, [], [f3274, f2686])).
fof(f3274, plain, (~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | spl68_436), inference(avatar_component_clause, [], [f3272])).
fof(f3272, plain, (spl68_436 <=> (h3(op1(e10, e12)) = op2(h3(e10), h3(e12)))), introduced(avatar_definition, [new_symbols(naming, [spl68_436])])).
fof(f6797, plain, (~ spl68_118 | ~ spl68_171 | ~ spl68_337 | ~ spl68_341 | ~ spl68_345 | spl68_435), inference(avatar_contradiction_clause, [], [f6796])).
fof(f6796, plain, ($false | (~ spl68_118 | ~ spl68_171 | ~ spl68_337 | ~ spl68_341 | ~ spl68_345 | spl68_435)), inference(subsumption_resolution, [], [f6795, f1655])).
fof(f6795, plain, (~ (e20 = op2(e23, e21)) | (~ spl68_118 | ~ spl68_337 | ~ spl68_341 | ~ spl68_345 | spl68_435)), inference(forward_demodulation, [], [f6794, f2686])).
fof(f6794, plain, (~ (op2(e23, e21) = h3(e12)) | (~ spl68_118 | ~ spl68_337 | ~ spl68_341 | spl68_435)), inference(forward_demodulation, [], [f6793, f1382])).
fof(f6793, plain, (~ (op2(e23, e21) = h3(op1(e10, e11))) | (~ spl68_337 | ~ spl68_341 | spl68_435)), inference(forward_demodulation, [], [f6792, f2645])).
fof(f6792, plain, (~ (h3(op1(e10, e11)) = op2(h3(e10), e21)) | (~ spl68_341 | spl68_435)), inference(forward_demodulation, [], [f3270, f2666])).
fof(f3270, plain, (~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | spl68_435), inference(avatar_component_clause, [], [f3268])).
fof(f3268, plain, (spl68_435 <=> (h3(op1(e10, e11)) = op2(h3(e10), h3(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl68_435])])).
fof(f6771, plain, (~ spl68_110 | ~ spl68_158 | ~ spl68_337 | spl68_437 | ~ spl68_450), inference(avatar_contradiction_clause, [], [f6770])).
fof(f6770, plain, ($false | (~ spl68_110 | ~ spl68_158 | ~ spl68_337 | spl68_437 | ~ spl68_450)), inference(subsumption_resolution, [], [f6769, f1600])).
fof(f6769, plain, (~ (e22 = op2(e23, e24)) | (~ spl68_110 | ~ spl68_337 | spl68_437 | ~ spl68_450)), inference(forward_demodulation, [], [f6768, f748])).
fof(f6768, plain, (~ (op2(e23, e24) = h3(e14)) | (~ spl68_110 | ~ spl68_337 | spl68_437 | ~ spl68_450)), inference(forward_demodulation, [], [f6767, f1348])).
fof(f6767, plain, (~ (op2(e23, e24) = h3(op1(e10, e13))) | (~ spl68_337 | spl68_437 | ~ spl68_450)), inference(forward_demodulation, [], [f6766, f2645])).
fof(f6766, plain, (~ (h3(op1(e10, e13)) = op2(h3(e10), e24)) | (spl68_437 | ~ spl68_450)), inference(forward_demodulation, [], [f3278, f3329])).
fof(f3278, plain, (~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | spl68_437), inference(avatar_component_clause, [], [f3276])).
fof(f6744, plain, (~ spl68_50 | ~ spl68_158 | ~ spl68_349 | ~ spl68_475 | ~ spl68_491), inference(avatar_contradiction_clause, [], [f6743])).
fof(f6743, plain, ($false | (~ spl68_50 | ~ spl68_158 | ~ spl68_349 | ~ spl68_475 | ~ spl68_491)), inference(subsumption_resolution, [], [f6742, f549])).
fof(f6742, plain, ((e21 = e22) | (~ spl68_50 | ~ spl68_158 | ~ spl68_349 | ~ spl68_475 | ~ spl68_491)), inference(forward_demodulation, [], [f6741, f743])).
fof(f6741, plain, ((e22 = h2(e14)) | (~ spl68_50 | ~ spl68_158 | ~ spl68_349 | ~ spl68_475 | ~ spl68_491)), inference(backward_demodulation, [], [f6709, f1096])).
fof(f6709, plain, ((e22 = h2(op1(e13, e10))) | (~ spl68_158 | ~ spl68_349 | ~ spl68_475 | ~ spl68_491)), inference(forward_demodulation, [], [f6635, f1600])).
fof(f6635, plain, ((op2(e23, e24) = h2(op1(e13, e10))) | (~ spl68_349 | ~ spl68_475 | ~ spl68_491)), inference(forward_demodulation, [], [f6634, f2706])).
fof(f2706, plain, ((e23 = h2(e13)) | ~ spl68_349), inference(avatar_component_clause, [], [f2705])).
fof(f2705, plain, (spl68_349 <=> (e23 = h2(e13))), introduced(avatar_definition, [new_symbols(naming, [spl68_349])])).
fof(f6634, plain, ((h2(op1(e13, e10)) = op2(h2(e13), e24)) | (~ spl68_475 | ~ spl68_491)), inference(forward_demodulation, [], [f3491, f3600])).
fof(f3600, plain, ((e24 = h2(e10)) | ~ spl68_491), inference(avatar_component_clause, [], [f3599])).
fof(f3599, plain, (spl68_491 <=> (e24 = h2(e10))), introduced(avatar_definition, [new_symbols(naming, [spl68_491])])).
fof(f3491, plain, ((h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ spl68_475), inference(avatar_component_clause, [], [f3490])).
fof(f6733, plain, (~ spl68_122 | ~ spl68_162 | ~ spl68_337 | ~ spl68_341 | spl68_434), inference(avatar_contradiction_clause, [], [f6732])).
fof(f6732, plain, ($false | (~ spl68_122 | ~ spl68_162 | ~ spl68_337 | ~ spl68_341 | spl68_434)), inference(subsumption_resolution, [], [f6728, f2666])).
fof(f6728, plain, (~ (e21 = h3(e11)) | (~ spl68_122 | ~ spl68_162 | ~ spl68_337 | spl68_434)), inference(backward_demodulation, [], [f6719, f1399])).
fof(f6719, plain, (~ (e21 = h3(op1(e10, e10))) | (~ spl68_162 | ~ spl68_337 | spl68_434)), inference(forward_demodulation, [], [f6718, f1617])).
fof(f6718, plain, (~ (op2(e23, e23) = h3(op1(e10, e10))) | (~ spl68_337 | spl68_434)), inference(forward_demodulation, [], [f3266, f2645])).
fof(f3266, plain, (~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10))) | spl68_434), inference(avatar_component_clause, [], [f3264])).
fof(f3264, plain, (spl68_434 <=> (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl68_434])])).
fof(f6730, plain, (~ spl68_47 | ~ spl68_122), inference(avatar_split_clause, [], [f6722, f1397, f1082])).
fof(f6722, plain, (~ (e11 = op1(e13, e10)) | ~ spl68_122), inference(backward_demodulation, [], [f337, f1399])).
fof(f337, plain, ~ (op1(e10, e10) = op1(e13, e10)), inference(cnf_transformation, [], [f7])).
fof(f6708, plain, (~ spl68_125 | ~ spl68_132 | spl68_463 | ~ spl68_491), inference(avatar_contradiction_clause, [], [f6707])).
fof(f6707, plain, ($false | (~ spl68_125 | ~ spl68_132 | spl68_463 | ~ spl68_491)), inference(subsumption_resolution, [], [f6706, f1491])).
fof(f6706, plain, (~ (e21 = op2(e24, e24)) | (~ spl68_125 | spl68_463 | ~ spl68_491)), inference(forward_demodulation, [], [f6705, f743])).
fof(f6705, plain, (~ (op2(e24, e24) = h2(e14)) | (~ spl68_125 | spl68_463 | ~ spl68_491)), inference(forward_demodulation, [], [f6704, f1411])).
fof(f6704, plain, (~ (op2(e24, e24) = h2(op1(e10, e10))) | (spl68_463 | ~ spl68_491)), inference(forward_demodulation, [], [f3444, f3600])).
fof(f6700, plain, (~ spl68_76 | spl68_487), inference(avatar_contradiction_clause, [], [f6699])).
fof(f6699, plain, ($false | (~ spl68_76 | spl68_487)), inference(trivial_inequality_removal, [], [f6698])).
fof(f6698, plain, (~ (h2(e10) = h2(e10)) | (~ spl68_76 | spl68_487)), inference(forward_demodulation, [], [f3540, f1206])).
fof(f3540, plain, (~ (h2(e10) = h2(op1(e11, e14))) | spl68_487), inference(avatar_component_clause, [], [f3538])).
fof(f3538, plain, (spl68_487 <=> (h2(e10) = h2(op1(e11, e14)))), introduced(avatar_definition, [new_symbols(naming, [spl68_487])])).
fof(f6697, plain, (~ spl68_55 | ~ spl68_247 | ~ spl68_360 | spl68_486), inference(avatar_contradiction_clause, [], [f6696])).
fof(f6696, plain, ($false | (~ spl68_55 | ~ spl68_247 | ~ spl68_360 | spl68_486)), inference(subsumption_resolution, [], [f6695, f1974])).
fof(f6695, plain, (~ (e21 = op2(e20, e21)) | (~ spl68_55 | ~ spl68_360 | spl68_486)), inference(forward_demodulation, [], [f6694, f743])).
fof(f6694, plain, (~ (op2(e20, e21) = h2(e14)) | (~ spl68_55 | ~ spl68_360 | spl68_486)), inference(forward_demodulation, [], [f6693, f1117])).
fof(f6693, plain, (~ (op2(e20, e21) = h2(op1(e12, e14))) | (~ spl68_360 | spl68_486)), inference(forward_demodulation, [], [f3536, f2762])).
fof(f2762, plain, ((e20 = h2(e12)) | ~ spl68_360), inference(avatar_component_clause, [], [f2761])).
fof(f2761, plain, (spl68_360 <=> (e20 = h2(e12))), introduced(avatar_definition, [new_symbols(naming, [spl68_360])])).
fof(f6692, plain, (~ spl68_28 | spl68_485), inference(avatar_contradiction_clause, [], [f6691])).
fof(f6691, plain, ($false | (~ spl68_28 | spl68_485)), inference(trivial_inequality_removal, [], [f6690])).
fof(f6690, plain, (~ (h2(e12) = h2(e12)) | (~ spl68_28 | spl68_485)), inference(forward_demodulation, [], [f3532, f1004])).
fof(f3532, plain, (~ (h2(e12) = h2(op1(e13, e14))) | spl68_485), inference(avatar_component_clause, [], [f3530])).
fof(f3530, plain, (spl68_485 <=> (h2(e12) = h2(op1(e13, e14)))), introduced(avatar_definition, [new_symbols(naming, [spl68_485])])).
fof(f6689, plain, (~ spl68_19 | spl68_483), inference(avatar_contradiction_clause, [], [f6688])).
fof(f6688, plain, ($false | (~ spl68_19 | spl68_483)), inference(trivial_inequality_removal, [], [f6687])).
fof(f6687, plain, (~ (h2(e13) = h2(e13)) | (~ spl68_19 | spl68_483)), inference(forward_demodulation, [], [f3524, f966])).
fof(f3524, plain, (~ (h2(e13) = h2(op1(e14, e11))) | spl68_483), inference(avatar_component_clause, [], [f3522])).
fof(f3522, plain, (spl68_483 <=> (h2(e13) = h2(op1(e14, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl68_483])])).
fof(f6680, plain, (~ spl68_23 | ~ spl68_206 | ~ spl68_360 | spl68_484 | ~ spl68_491), inference(avatar_contradiction_clause, [], [f6679])).
fof(f6679, plain, ($false | (~ spl68_23 | ~ spl68_206 | ~ spl68_360 | spl68_484 | ~ spl68_491)), inference(subsumption_resolution, [], [f6678, f1802])).
fof(f6678, plain, (~ (e20 = op2(e21, e24)) | (~ spl68_23 | ~ spl68_360 | spl68_484 | ~ spl68_491)), inference(forward_demodulation, [], [f6677, f2762])).
fof(f6677, plain, (~ (op2(e21, e24) = h2(e12)) | (~ spl68_23 | spl68_484 | ~ spl68_491)), inference(forward_demodulation, [], [f6676, f983])).
fof(f6676, plain, (~ (op2(e21, e24) = h2(op1(e14, e10))) | (spl68_484 | ~ spl68_491)), inference(forward_demodulation, [], [f3528, f3600])).
fof(f6675, plain, (~ spl68_15 | ~ spl68_227 | ~ spl68_360 | spl68_482), inference(avatar_contradiction_clause, [], [f6674])).
fof(f6674, plain, ($false | (~ spl68_15 | ~ spl68_227 | ~ spl68_360 | spl68_482)), inference(subsumption_resolution, [], [f6673, f1890])).
fof(f6673, plain, (~ (e21 = op2(e21, e20)) | (~ spl68_15 | ~ spl68_360 | spl68_482)), inference(forward_demodulation, [], [f6672, f743])).
fof(f6672, plain, (~ (op2(e21, e20) = h2(e14)) | (~ spl68_15 | ~ spl68_360 | spl68_482)), inference(forward_demodulation, [], [f6671, f949])).
fof(f6671, plain, (~ (op2(e21, e20) = h2(op1(e14, e12))) | (~ spl68_360 | spl68_482)), inference(forward_demodulation, [], [f3520, f2762])).
fof(f6670, plain, (~ spl68_2 | spl68_480), inference(avatar_contradiction_clause, [], [f6669])).
fof(f6669, plain, ($false | (~ spl68_2 | spl68_480)), inference(trivial_inequality_removal, [], [f6668])).
fof(f6668, plain, (~ (h2(e11) = h2(e11)) | (~ spl68_2 | spl68_480)), inference(forward_demodulation, [], [f3512, f895])).
fof(f3512, plain, (~ (h2(e11) = h2(op1(e14, e14))) | spl68_480), inference(avatar_component_clause, [], [f3510])).
fof(f3510, plain, (spl68_480 <=> (h2(e11) = h2(op1(e14, e14)))), introduced(avatar_definition, [new_symbols(naming, [spl68_480])])).
fof(f6661, plain, (~ spl68_6 | ~ spl68_215 | ~ spl68_349 | spl68_481 | ~ spl68_491), inference(avatar_contradiction_clause, [], [f6660])).
fof(f6660, plain, ($false | (~ spl68_6 | ~ spl68_215 | ~ spl68_349 | spl68_481 | ~ spl68_491)), inference(subsumption_resolution, [], [f6659, f1839])).
fof(f6659, plain, (~ (e24 = op2(e21, e23)) | (~ spl68_6 | ~ spl68_349 | spl68_481 | ~ spl68_491)), inference(forward_demodulation, [], [f6658, f3600])).
fof(f6658, plain, (~ (op2(e21, e23) = h2(e10)) | (~ spl68_6 | ~ spl68_349 | spl68_481)), inference(forward_demodulation, [], [f6657, f912])).
fof(f6657, plain, (~ (op2(e21, e23) = h2(op1(e14, e13))) | (~ spl68_349 | spl68_481)), inference(forward_demodulation, [], [f3516, f2706])).
fof(f6652, plain, (~ spl68_35 | ~ spl68_162 | ~ spl68_349 | spl68_478), inference(avatar_contradiction_clause, [], [f6651])).
fof(f6651, plain, ($false | (~ spl68_35 | ~ spl68_162 | ~ spl68_349 | spl68_478)), inference(subsumption_resolution, [], [f6650, f1617])).
fof(f6650, plain, (~ (e21 = op2(e23, e23)) | (~ spl68_35 | ~ spl68_349 | spl68_478)), inference(forward_demodulation, [], [f6649, f743])).
fof(f6649, plain, (~ (op2(e23, e23) = h2(e14)) | (~ spl68_35 | ~ spl68_349 | spl68_478)), inference(forward_demodulation, [], [f6648, f1033])).
fof(f6648, plain, (~ (op2(e23, e23) = h2(op1(e13, e13))) | (~ spl68_349 | spl68_478)), inference(forward_demodulation, [], [f3504, f2706])).
fof(f6643, plain, (~ spl68_39 | ~ spl68_179 | ~ spl68_349 | ~ spl68_360 | spl68_477), inference(avatar_contradiction_clause, [], [f6642])).
fof(f6642, plain, ($false | (~ spl68_39 | ~ spl68_179 | ~ spl68_349 | ~ spl68_360 | spl68_477)), inference(subsumption_resolution, [], [f6641, f1688])).
fof(f6641, plain, (~ (e23 = op2(e23, e20)) | (~ spl68_39 | ~ spl68_349 | ~ spl68_360 | spl68_477)), inference(forward_demodulation, [], [f6640, f2706])).
fof(f6640, plain, (~ (op2(e23, e20) = h2(e13)) | (~ spl68_39 | ~ spl68_349 | ~ spl68_360 | spl68_477)), inference(forward_demodulation, [], [f6639, f1050])).
fof(f6639, plain, (~ (op2(e23, e20) = h2(op1(e13, e12))) | (~ spl68_349 | ~ spl68_360 | spl68_477)), inference(forward_demodulation, [], [f6638, f2706])).
fof(f6638, plain, (~ (h2(op1(e13, e12)) = op2(h2(e13), e20)) | (~ spl68_360 | spl68_477)), inference(forward_demodulation, [], [f3500, f2762])).
fof(f6633, plain, (~ spl68_41 | ~ spl68_170 | ~ spl68_349 | ~ spl68_356 | spl68_476 | ~ spl68_491), inference(avatar_contradiction_clause, [], [f6632])).
fof(f6632, plain, ($false | (~ spl68_41 | ~ spl68_170 | ~ spl68_349 | ~ spl68_356 | spl68_476 | ~ spl68_491)), inference(subsumption_resolution, [], [f6631, f1650])).
fof(f6631, plain, (~ (e24 = op2(e23, e22)) | (~ spl68_41 | ~ spl68_349 | ~ spl68_356 | spl68_476 | ~ spl68_491)), inference(forward_demodulation, [], [f6630, f3600])).
fof(f6630, plain, (~ (op2(e23, e22) = h2(e10)) | (~ spl68_41 | ~ spl68_349 | ~ spl68_356 | spl68_476)), inference(forward_demodulation, [], [f6629, f1059])).
fof(f6629, plain, (~ (op2(e23, e22) = h2(op1(e13, e11))) | (~ spl68_349 | ~ spl68_356 | spl68_476)), inference(forward_demodulation, [], [f6628, f2706])).
fof(f6628, plain, (~ (h2(op1(e13, e11)) = op2(h2(e13), e22)) | (~ spl68_356 | spl68_476)), inference(forward_demodulation, [], [f3496, f2741])).
fof(f2741, plain, ((e22 = h2(e11)) | ~ spl68_356), inference(avatar_component_clause, [], [f2740])).
fof(f2740, plain, (spl68_356 <=> (e22 = h2(e11))), introduced(avatar_definition, [new_symbols(naming, [spl68_356])])).
fof(f6623, plain, (~ spl68_47 | ~ spl68_158 | ~ spl68_349 | ~ spl68_356 | spl68_475 | ~ spl68_491), inference(avatar_contradiction_clause, [], [f6622])).
fof(f6622, plain, ($false | (~ spl68_47 | ~ spl68_158 | ~ spl68_349 | ~ spl68_356 | spl68_475 | ~ spl68_491)), inference(subsumption_resolution, [], [f6621, f1600])).
fof(f6621, plain, (~ (e22 = op2(e23, e24)) | (~ spl68_47 | ~ spl68_349 | ~ spl68_356 | spl68_475 | ~ spl68_491)), inference(forward_demodulation, [], [f6620, f2741])).
fof(f6620, plain, (~ (op2(e23, e24) = h2(e11)) | (~ spl68_47 | ~ spl68_349 | spl68_475 | ~ spl68_491)), inference(forward_demodulation, [], [f6619, f1084])).
fof(f6619, plain, (~ (op2(e23, e24) = h2(op1(e13, e10))) | (~ spl68_349 | spl68_475 | ~ spl68_491)), inference(forward_demodulation, [], [f6618, f2706])).
fof(f6618, plain, (~ (h2(op1(e13, e10)) = op2(h2(e13), e24)) | (spl68_475 | ~ spl68_491)), inference(forward_demodulation, [], [f3492, f3600])).
fof(f6614, plain, (~ spl68_59 | ~ spl68_239 | ~ spl68_349 | ~ spl68_360 | spl68_474), inference(avatar_contradiction_clause, [], [f6613])).
fof(f6613, plain, ($false | (~ spl68_59 | ~ spl68_239 | ~ spl68_349 | ~ spl68_360 | spl68_474)), inference(subsumption_resolution, [], [f6612, f1940])).
fof(f6612, plain, (~ (e23 = op2(e20, e23)) | (~ spl68_59 | ~ spl68_349 | ~ spl68_360 | spl68_474)), inference(forward_demodulation, [], [f6611, f2706])).
fof(f6611, plain, (~ (op2(e20, e23) = h2(e13)) | (~ spl68_59 | ~ spl68_349 | ~ spl68_360 | spl68_474)), inference(forward_demodulation, [], [f6610, f1134])).
fof(f6610, plain, (~ (op2(e20, e23) = h2(op1(e12, e13))) | (~ spl68_349 | ~ spl68_360 | spl68_474)), inference(forward_demodulation, [], [f6609, f2762])).
fof(f6609, plain, (~ (h2(op1(e12, e13)) = op2(h2(e12), e23)) | (~ spl68_349 | spl68_474)), inference(forward_demodulation, [], [f3488, f2706])).
fof(f6604, plain, (~ spl68_63 | ~ spl68_251 | ~ spl68_360 | spl68_473), inference(avatar_contradiction_clause, [], [f6603])).
fof(f6603, plain, ($false | (~ spl68_63 | ~ spl68_251 | ~ spl68_360 | spl68_473)), inference(subsumption_resolution, [], [f6602, f1991])).
fof(f6602, plain, (~ (e20 = op2(e20, e20)) | (~ spl68_63 | ~ spl68_360 | spl68_473)), inference(forward_demodulation, [], [f6601, f2762])).
fof(f6601, plain, (~ (op2(e20, e20) = h2(e12)) | (~ spl68_63 | ~ spl68_360 | spl68_473)), inference(forward_demodulation, [], [f6600, f1151])).
fof(f6600, plain, (~ (op2(e20, e20) = h2(op1(e12, e12))) | (~ spl68_360 | spl68_473)), inference(forward_demodulation, [], [f3484, f2762])).
fof(f6595, plain, (~ spl68_67 | ~ spl68_243 | ~ spl68_356 | ~ spl68_360 | spl68_472), inference(avatar_contradiction_clause, [], [f6594])).
fof(f6594, plain, ($false | (~ spl68_67 | ~ spl68_243 | ~ spl68_356 | ~ spl68_360 | spl68_472)), inference(subsumption_resolution, [], [f6593, f1957])).
fof(f6593, plain, (~ (e22 = op2(e20, e22)) | (~ spl68_67 | ~ spl68_356 | ~ spl68_360 | spl68_472)), inference(forward_demodulation, [], [f6592, f2741])).
fof(f6592, plain, (~ (op2(e20, e22) = h2(e11)) | (~ spl68_67 | ~ spl68_356 | ~ spl68_360 | spl68_472)), inference(forward_demodulation, [], [f6591, f1168])).
fof(f6591, plain, (~ (op2(e20, e22) = h2(op1(e12, e11))) | (~ spl68_356 | ~ spl68_360 | spl68_472)), inference(forward_demodulation, [], [f6590, f2762])).
fof(f6590, plain, (~ (h2(op1(e12, e11)) = op2(h2(e12), e22)) | (~ spl68_356 | spl68_472)), inference(forward_demodulation, [], [f3480, f2741])).
fof(f6585, plain, (~ spl68_71 | ~ spl68_235 | ~ spl68_360 | spl68_471 | ~ spl68_491), inference(avatar_contradiction_clause, [], [f6584])).
fof(f6584, plain, ($false | (~ spl68_71 | ~ spl68_235 | ~ spl68_360 | spl68_471 | ~ spl68_491)), inference(subsumption_resolution, [], [f6583, f1923])).
fof(f6583, plain, (~ (e24 = op2(e20, e24)) | (~ spl68_71 | ~ spl68_360 | spl68_471 | ~ spl68_491)), inference(forward_demodulation, [], [f6582, f3600])).
fof(f6582, plain, (~ (op2(e20, e24) = h2(e10)) | (~ spl68_71 | ~ spl68_360 | spl68_471 | ~ spl68_491)), inference(forward_demodulation, [], [f6581, f1185])).
fof(f6581, plain, (~ (op2(e20, e24) = h2(op1(e12, e10))) | (~ spl68_360 | spl68_471 | ~ spl68_491)), inference(forward_demodulation, [], [f6580, f2762])).
fof(f6580, plain, (~ (h2(op1(e12, e10)) = op2(h2(e12), e24)) | (spl68_471 | ~ spl68_491)), inference(forward_demodulation, [], [f3476, f3600])).
fof(f6575, plain, (~ spl68_83 | ~ spl68_186 | ~ spl68_349 | ~ spl68_356 | ~ spl68_360 | spl68_470), inference(avatar_contradiction_clause, [], [f6574])).
fof(f6574, plain, ($false | (~ spl68_83 | ~ spl68_186 | ~ spl68_349 | ~ spl68_356 | ~ spl68_360 | spl68_470)), inference(subsumption_resolution, [], [f6573, f1718])).
fof(f6573, plain, (~ (e20 = op2(e22, e23)) | (~ spl68_83 | ~ spl68_349 | ~ spl68_356 | ~ spl68_360 | spl68_470)), inference(forward_demodulation, [], [f6572, f2762])).
fof(f6572, plain, (~ (op2(e22, e23) = h2(e12)) | (~ spl68_83 | ~ spl68_349 | ~ spl68_356 | spl68_470)), inference(forward_demodulation, [], [f6571, f1235])).
fof(f6571, plain, (~ (op2(e22, e23) = h2(op1(e11, e13))) | (~ spl68_349 | ~ spl68_356 | spl68_470)), inference(forward_demodulation, [], [f6570, f2741])).
fof(f6570, plain, (~ (h2(op1(e11, e13)) = op2(h2(e11), e23)) | (~ spl68_349 | spl68_470)), inference(forward_demodulation, [], [f3472, f2706])).
fof(f6566, plain, (~ spl68_87 | ~ spl68_203 | ~ spl68_356 | ~ spl68_360 | spl68_469), inference(avatar_contradiction_clause, [], [f6565])).
fof(f6565, plain, ($false | (~ spl68_87 | ~ spl68_203 | ~ spl68_356 | ~ spl68_360 | spl68_469)), inference(subsumption_resolution, [], [f6564, f1789])).
fof(f6564, plain, (~ (e22 = op2(e22, e20)) | (~ spl68_87 | ~ spl68_356 | ~ spl68_360 | spl68_469)), inference(forward_demodulation, [], [f6563, f2741])).
fof(f6563, plain, (~ (op2(e22, e20) = h2(e11)) | (~ spl68_87 | ~ spl68_356 | ~ spl68_360 | spl68_469)), inference(forward_demodulation, [], [f6562, f1252])).
fof(f6562, plain, (~ (op2(e22, e20) = h2(op1(e11, e12))) | (~ spl68_356 | ~ spl68_360 | spl68_469)), inference(forward_demodulation, [], [f6561, f2741])).
fof(f6561, plain, (~ (h2(op1(e11, e12)) = op2(h2(e11), e20)) | (~ spl68_360 | spl68_469)), inference(forward_demodulation, [], [f3468, f2762])).
fof(f6556, plain, (~ spl68_95 | ~ spl68_192 | ~ spl68_356 | spl68_468), inference(avatar_contradiction_clause, [], [f6555])).
fof(f6555, plain, ($false | (~ spl68_95 | ~ spl68_192 | ~ spl68_356 | spl68_468)), inference(subsumption_resolution, [], [f6554, f1743])).
fof(f6554, plain, (~ (e21 = op2(e22, e22)) | (~ spl68_95 | ~ spl68_356 | spl68_468)), inference(forward_demodulation, [], [f6553, f743])).
fof(f6553, plain, (~ (op2(e22, e22) = h2(e14)) | (~ spl68_95 | ~ spl68_356 | spl68_468)), inference(forward_demodulation, [], [f6552, f1285])).
fof(f6552, plain, (~ (op2(e22, e22) = h2(op1(e11, e11))) | (~ spl68_356 | spl68_468)), inference(forward_demodulation, [], [f3464, f2741])).
fof(f6547, plain, (~ spl68_99 | ~ spl68_184 | ~ spl68_349 | ~ spl68_356 | spl68_467 | ~ spl68_491), inference(avatar_contradiction_clause, [], [f6546])).
fof(f6546, plain, ($false | (~ spl68_99 | ~ spl68_184 | ~ spl68_349 | ~ spl68_356 | spl68_467 | ~ spl68_491)), inference(subsumption_resolution, [], [f6545, f1709])).
fof(f6545, plain, (~ (e23 = op2(e22, e24)) | (~ spl68_99 | ~ spl68_349 | ~ spl68_356 | spl68_467 | ~ spl68_491)), inference(forward_demodulation, [], [f6544, f2706])).
fof(f6544, plain, (~ (op2(e22, e24) = h2(e13)) | (~ spl68_99 | ~ spl68_356 | spl68_467 | ~ spl68_491)), inference(forward_demodulation, [], [f6543, f1302])).
fof(f6543, plain, (~ (op2(e22, e24) = h2(op1(e11, e10))) | (~ spl68_356 | spl68_467 | ~ spl68_491)), inference(forward_demodulation, [], [f6542, f2741])).
fof(f6542, plain, (~ (h2(op1(e11, e10)) = op2(h2(e11), e24)) | (spl68_467 | ~ spl68_491)), inference(forward_demodulation, [], [f3460, f3600])).
fof(f6537, plain, (~ spl68_107 | ~ spl68_138 | ~ spl68_349 | ~ spl68_356 | spl68_466 | ~ spl68_491), inference(avatar_contradiction_clause, [], [f6536])).
fof(f6536, plain, ($false | (~ spl68_107 | ~ spl68_138 | ~ spl68_349 | ~ spl68_356 | spl68_466 | ~ spl68_491)), inference(subsumption_resolution, [], [f6535, f1516])).
fof(f6535, plain, (~ (e22 = op2(e24, e23)) | (~ spl68_107 | ~ spl68_349 | ~ spl68_356 | spl68_466 | ~ spl68_491)), inference(forward_demodulation, [], [f6534, f2741])).
fof(f6534, plain, (~ (op2(e24, e23) = h2(e11)) | (~ spl68_107 | ~ spl68_349 | spl68_466 | ~ spl68_491)), inference(forward_demodulation, [], [f6533, f1336])).
fof(f6533, plain, (~ (op2(e24, e23) = h2(op1(e10, e13))) | (~ spl68_349 | spl68_466 | ~ spl68_491)), inference(forward_demodulation, [], [f6532, f3600])).
fof(f6532, plain, (~ (h2(op1(e10, e13)) = op2(h2(e10), e23)) | (~ spl68_349 | spl68_466)), inference(forward_demodulation, [], [f3456, f2706])).
fof(f6527, plain, (~ spl68_111 | ~ spl68_155 | ~ spl68_360 | spl68_465 | ~ spl68_491), inference(avatar_contradiction_clause, [], [f6526])).
fof(f6526, plain, ($false | (~ spl68_111 | ~ spl68_155 | ~ spl68_360 | spl68_465 | ~ spl68_491)), inference(subsumption_resolution, [], [f6525, f1587])).
fof(f6525, plain, (~ (e24 = op2(e24, e20)) | (~ spl68_111 | ~ spl68_360 | spl68_465 | ~ spl68_491)), inference(forward_demodulation, [], [f6524, f3600])).
fof(f6524, plain, (~ (op2(e24, e20) = h2(e10)) | (~ spl68_111 | ~ spl68_360 | spl68_465 | ~ spl68_491)), inference(forward_demodulation, [], [f6523, f1353])).
fof(f6523, plain, (~ (op2(e24, e20) = h2(op1(e10, e12))) | (~ spl68_360 | spl68_465 | ~ spl68_491)), inference(forward_demodulation, [], [f6522, f3600])).
fof(f6522, plain, (~ (h2(op1(e10, e12)) = op2(h2(e10), e20)) | (~ spl68_360 | spl68_465)), inference(forward_demodulation, [], [f3452, f2762])).
fof(f6521, plain, (~ spl68_360 | ~ spl68_118 | ~ spl68_141 | ~ spl68_356 | spl68_464 | ~ spl68_491), inference(avatar_split_clause, [], [f6520, f3599, f3446, f2740, f1527, f1380, f2761])).
fof(f6520, plain, (~ (e20 = h2(e12)) | (~ spl68_118 | ~ spl68_141 | ~ spl68_356 | spl68_464 | ~ spl68_491)), inference(forward_demodulation, [], [f6519, f1529])).
fof(f6519, plain, (~ (op2(e24, e22) = h2(e12)) | (~ spl68_118 | ~ spl68_356 | spl68_464 | ~ spl68_491)), inference(forward_demodulation, [], [f6518, f1382])).
fof(f6518, plain, (~ (op2(e24, e22) = h2(op1(e10, e11))) | (~ spl68_356 | spl68_464 | ~ spl68_491)), inference(forward_demodulation, [], [f6517, f3600])).
fof(f6517, plain, (~ (h2(op1(e10, e11)) = op2(h2(e10), e22)) | (~ spl68_356 | spl68_464)), inference(forward_demodulation, [], [f3448, f2741])).
fof(f6510, plain, (spl68_360 | ~ spl68_171 | ~ spl68_349), inference(avatar_split_clause, [], [f6509, f2705, f1653, f2761])).
fof(f6509, plain, ((e20 = h2(e12)) | (~ spl68_171 | ~ spl68_349)), inference(forward_demodulation, [], [f6508, f1655])).
fof(f6508, plain, ((op2(e23, e21) = h2(e12)) | ~ spl68_349), inference(backward_demodulation, [], [f2439, f2706])).
fof(f6507, plain, (~ spl68_349 | ~ spl68_104 | ~ spl68_149 | spl68_488 | ~ spl68_491), inference(avatar_split_clause, [], [f6506, f3599, f3542, f1560, f1321, f2705])).
fof(f3542, plain, (spl68_488 <=> (h2(op1(e10, e14)) = op2(h2(e10), e21))), introduced(avatar_definition, [new_symbols(naming, [spl68_488])])).
fof(f6506, plain, (~ (e23 = h2(e13)) | (~ spl68_104 | ~ spl68_149 | spl68_488 | ~ spl68_491)), inference(forward_demodulation, [], [f6504, f1562])).
fof(f6504, plain, (~ (op2(e24, e21) = h2(e13)) | (~ spl68_104 | spl68_488 | ~ spl68_491)), inference(backward_demodulation, [], [f6479, f3600])).
fof(f6479, plain, (~ (h2(e13) = op2(h2(e10), e21)) | (~ spl68_104 | spl68_488)), inference(forward_demodulation, [], [f3544, f1323])).
fof(f3544, plain, (~ (h2(op1(e10, e14)) = op2(h2(e10), e21)) | spl68_488), inference(avatar_component_clause, [], [f3542])).
fof(f6503, plain, (spl68_337 | ~ spl68_219 | ~ spl68_341), inference(avatar_split_clause, [], [f6485, f2665, f1854, f2644])).
fof(f6485, plain, ((e23 = h3(e10)) | (~ spl68_219 | ~ spl68_341)), inference(forward_demodulation, [], [f6484, f1856])).
fof(f6484, plain, ((op2(e21, e22) = h3(e10)) | ~ spl68_341), inference(forward_demodulation, [], [f2460, f2666])).
fof(f2460, plain, (h3(e10) = op2(h3(e11), e22)), inference(forward_demodulation, [], [f749, f750])).
fof(f749, plain, (h3(e10) = op2(op2(e22, e22), e22)), inference(cnf_transformation, [], [f18])).
fof(f6492, plain, (spl68_345 | ~ spl68_141 | ~ spl68_450), inference(avatar_split_clause, [], [f6491, f3328, f1527, f2685])).
fof(f6491, plain, ((e20 = h3(e12)) | (~ spl68_141 | ~ spl68_450)), inference(forward_demodulation, [], [f6490, f1529])).
fof(f6490, plain, ((op2(e24, e22) = h3(e12)) | ~ spl68_450), inference(backward_demodulation, [], [f2450, f3329])).
fof(f2450, plain, (h3(e12) = op2(h3(e13), e22)), inference(forward_demodulation, [], [f751, f752])).
fof(f752, plain, (op2(e22, op2(e22, e22)) = h3(e13)), inference(cnf_transformation, [], [f18])).
fof(f751, plain, (h3(e12) = op2(op2(e22, op2(e22, e22)), e22)), inference(cnf_transformation, [], [f18])).
fof(f6483, plain, (spl68_349 | ~ spl68_219 | ~ spl68_356), inference(avatar_split_clause, [], [f6482, f2740, f1854, f2705])).
fof(f6482, plain, ((e23 = h2(e13)) | (~ spl68_219 | ~ spl68_356)), inference(forward_demodulation, [], [f6420, f1856])).
fof(f6420, plain, ((op2(e21, e22) = h2(e13)) | ~ spl68_356), inference(forward_demodulation, [], [f2448, f2741])).
fof(f6481, plain, (spl68_491 | ~ spl68_200 | ~ spl68_356), inference(avatar_split_clause, [], [f6480, f2740, f1774, f3599])).
fof(f6480, plain, ((e24 = h2(e10)) | (~ spl68_200 | ~ spl68_356)), inference(forward_demodulation, [], [f6416, f1776])).
fof(f6416, plain, ((op2(e22, e21) = h2(e10)) | ~ spl68_356), inference(forward_demodulation, [], [f2449, f2741])).
fof(f6472, plain, (~ spl68_1 | ~ spl68_2), inference(avatar_contradiction_clause, [], [f6471])).
fof(f6471, plain, ($false | (~ spl68_1 | ~ spl68_2)), inference(subsumption_resolution, [], [f6470, f535])).
fof(f6470, plain, ((e10 = e11) | (~ spl68_1 | ~ spl68_2)), inference(backward_demodulation, [], [f895, f891])).
fof(f891, plain, ((e10 = op1(e14, e14)) | ~ spl68_1), inference(avatar_component_clause, [], [f889])).
fof(f889, plain, (spl68_1 <=> (e10 = op1(e14, e14))), introduced(avatar_definition, [new_symbols(naming, [spl68_1])])).
fof(f6459, plain, (~ spl68_223 | ~ spl68_224), inference(avatar_contradiction_clause, [], [f6458])).
fof(f6458, plain, ($false | (~ spl68_223 | ~ spl68_224)), inference(subsumption_resolution, [], [f6457, f552])).
fof(f552, plain, ~ (e22 = e23), inference(cnf_transformation, [], [f10])).
fof(f6457, plain, ((e22 = e23) | (~ spl68_223 | ~ spl68_224)), inference(forward_demodulation, [], [f1877, f1873])).
fof(f1877, plain, ((e23 = op2(e21, e21)) | ~ spl68_224), inference(avatar_component_clause, [], [f1875])).
fof(f1875, plain, (spl68_224 <=> (e23 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl68_224])])).
fof(f6452, plain, (spl68_450 | ~ spl68_200 | ~ spl68_341), inference(avatar_split_clause, [], [f6451, f2665, f1774, f3328])).
fof(f6451, plain, ((e24 = h3(e13)) | (~ spl68_200 | ~ spl68_341)), inference(forward_demodulation, [], [f6446, f1776])).
fof(f6446, plain, ((op2(e22, e21) = h3(e13)) | ~ spl68_341), inference(backward_demodulation, [], [f2459, f2666])).
fof(f2459, plain, (h3(e13) = op2(e22, h3(e11))), inference(backward_demodulation, [], [f752, f750])).
fof(f6432, plain, (~ spl68_104 | ~ spl68_171 | ~ spl68_352 | ~ spl68_359 | spl68_488), inference(avatar_contradiction_clause, [], [f6431])).
fof(f6431, plain, ($false | (~ spl68_104 | ~ spl68_171 | ~ spl68_352 | ~ spl68_359 | spl68_488)), inference(subsumption_resolution, [], [f6430, f1655])).
fof(f6430, plain, (~ (e20 = op2(e23, e21)) | (~ spl68_104 | ~ spl68_352 | ~ spl68_359 | spl68_488)), inference(forward_demodulation, [], [f6429, f2757])).
fof(f6429, plain, (~ (op2(e23, e21) = h2(e13)) | (~ spl68_104 | ~ spl68_352 | spl68_488)), inference(forward_demodulation, [], [f6428, f1323])).
fof(f6428, plain, (~ (op2(e23, e21) = h2(op1(e10, e14))) | (~ spl68_352 | spl68_488)), inference(forward_demodulation, [], [f3544, f2721])).
fof(f6414, plain, (spl68_341 | ~ spl68_192), inference(avatar_split_clause, [], [f6313, f1741, f2665])).
fof(f6313, plain, ((e21 = h3(e11)) | ~ spl68_192), inference(forward_demodulation, [], [f750, f1743])).
fof(f6397, plain, (~ spl68_14 | ~ spl68_19), inference(avatar_split_clause, [], [f6396, f964, f943])).
fof(f943, plain, (spl68_14 <=> (e13 = op1(e14, e12))), introduced(avatar_definition, [new_symbols(naming, [spl68_14])])).
fof(f6396, plain, (~ (e13 = op1(e14, e12)) | ~ spl68_19), inference(forward_demodulation, [], [f429, f966])).
fof(f429, plain, ~ (op1(e14, e11) = op1(e14, e12)), inference(cnf_transformation, [], [f7])).
fof(f6377, plain, (~ spl68_28 | ~ spl68_30), inference(avatar_contradiction_clause, [], [f6376])).
fof(f6376, plain, ($false | (~ spl68_28 | ~ spl68_30)), inference(subsumption_resolution, [], [f6375, f543])).
fof(f543, plain, ~ (e12 = e14), inference(cnf_transformation, [], [f9])).
fof(f6375, plain, ((e12 = e14) | (~ spl68_28 | ~ spl68_30)), inference(forward_demodulation, [], [f1012, f1004])).
fof(f1012, plain, ((e14 = op1(e13, e14)) | ~ spl68_30), inference(avatar_component_clause, [], [f1010])).
fof(f1010, plain, (spl68_30 <=> (e14 = op1(e13, e14))), introduced(avatar_definition, [new_symbols(naming, [spl68_30])])).
fof(f6364, plain, (~ spl68_125 | ~ spl68_162 | ~ spl68_352 | spl68_463), inference(avatar_contradiction_clause, [], [f6363])).
fof(f6363, plain, ($false | (~ spl68_125 | ~ spl68_162 | ~ spl68_352 | spl68_463)), inference(subsumption_resolution, [], [f6359, f743])).
fof(f6359, plain, (~ (e21 = h2(e14)) | (~ spl68_125 | ~ spl68_162 | ~ spl68_352 | spl68_463)), inference(backward_demodulation, [], [f6335, f1411])).
fof(f6335, plain, (~ (e21 = h2(op1(e10, e10))) | (~ spl68_162 | ~ spl68_352 | spl68_463)), inference(forward_demodulation, [], [f6334, f1617])).
fof(f6334, plain, (~ (op2(e23, e23) = h2(op1(e10, e10))) | (~ spl68_352 | spl68_463)), inference(backward_demodulation, [], [f3444, f2721])).
fof(f6360, plain, (~ spl68_50 | ~ spl68_125), inference(avatar_split_clause, [], [f6355, f1409, f1094])).
fof(f6355, plain, (~ (e14 = op1(e13, e10)) | ~ spl68_125), inference(backward_demodulation, [], [f337, f1411])).
fof(f6340, plain, (~ spl68_223 | ~ spl68_225), inference(avatar_contradiction_clause, [], [f6339])).
fof(f6339, plain, ($false | (~ spl68_223 | ~ spl68_225)), inference(subsumption_resolution, [], [f6338, f553])).
fof(f553, plain, ~ (e22 = e24), inference(cnf_transformation, [], [f10])).
fof(f6338, plain, ((e22 = e24) | (~ spl68_223 | ~ spl68_225)), inference(backward_demodulation, [], [f1881, f1873])).
fof(f6320, plain, (~ spl68_198 | ~ spl68_203), inference(avatar_split_clause, [], [f6319, f1787, f1766])).
fof(f6319, plain, (~ (e22 = op2(e22, e21)) | ~ spl68_203), inference(forward_demodulation, [], [f505, f1789])).
fof(f505, plain, ~ (op2(e22, e20) = op2(e22, e21)), inference(cnf_transformation, [], [f8])).
fof(f6217, plain, (~ spl68_167 | ~ spl68_170), inference(avatar_contradiction_clause, [], [f6216])).
fof(f6216, plain, ($false | (~ spl68_167 | ~ spl68_170)), inference(subsumption_resolution, [], [f6215, f551])).
fof(f6215, plain, ((e21 = e24) | (~ spl68_167 | ~ spl68_170)), inference(forward_demodulation, [], [f1650, f1638])).
fof(f1638, plain, ((e21 = op2(e23, e22)) | ~ spl68_167), inference(avatar_component_clause, [], [f1636])).
fof(f1636, plain, (spl68_167 <=> (e21 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl68_167])])).
fof(f6194, plain, (~ spl68_224 | ~ spl68_225), inference(avatar_contradiction_clause, [], [f6193])).
fof(f6193, plain, ($false | (~ spl68_224 | ~ spl68_225)), inference(subsumption_resolution, [], [f6192, f554])).
fof(f554, plain, ~ (e23 = e24), inference(cnf_transformation, [], [f10])).
fof(f6192, plain, ((e23 = e24) | (~ spl68_224 | ~ spl68_225)), inference(backward_demodulation, [], [f1881, f1877])).
fof(f6172, plain, (~ spl68_19 | spl68_385), inference(avatar_contradiction_clause, [], [f6171])).
fof(f6171, plain, ($false | (~ spl68_19 | spl68_385)), inference(subsumption_resolution, [], [f6170, f2472])).
fof(f6170, plain, (~ (e23 = h5(e13)) | (~ spl68_19 | spl68_385)), inference(forward_demodulation, [], [f2917, f966])).
fof(f2917, plain, (~ (e23 = h5(op1(e14, e11))) | spl68_385), inference(avatar_component_clause, [], [f2915])).
fof(f2915, plain, (spl68_385 <=> (e23 = h5(op1(e14, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl68_385])])).
fof(f6134, plain, (spl68_490 | ~ spl68_225), inference(avatar_split_clause, [], [f5945, f1879, f3580])).
fof(f5945, plain, ((e24 = h2(e11)) | ~ spl68_225), inference(backward_demodulation, [], [f745, f1881])).
fof(f6113, plain, (~ spl68_212 | ~ spl68_227), inference(avatar_split_clause, [], [f6112, f1888, f1825])).
fof(f1825, plain, (spl68_212 <=> (e21 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl68_212])])).
fof(f6112, plain, (~ (e21 = op2(e21, e23)) | ~ spl68_227), inference(forward_demodulation, [], [f497, f1890])).
fof(f497, plain, ~ (op2(e21, e20) = op2(e21, e23)), inference(cnf_transformation, [], [f8])).
fof(f6089, plain, (~ spl68_123 | ~ spl68_23), inference(avatar_split_clause, [], [f6088, f981, f1401])).
fof(f1401, plain, (spl68_123 <=> (op1(e10, e10) = e12)), introduced(avatar_definition, [new_symbols(naming, [spl68_123])])).
fof(f6088, plain, (~ (op1(e10, e10) = e12) | ~ spl68_23), inference(forward_demodulation, [], [f338, f983])).
fof(f338, plain, ~ (op1(e10, e10) = op1(e14, e10)), inference(cnf_transformation, [], [f7])).
fof(f6066, plain, (~ spl68_22 | ~ spl68_2), inference(avatar_split_clause, [], [f4776, f893, f977])).
fof(f977, plain, (spl68_22 <=> (e11 = op1(e14, e10))), introduced(avatar_definition, [new_symbols(naming, [spl68_22])])).
fof(f4776, plain, (~ (e11 = op1(e14, e10)) | ~ spl68_2), inference(forward_demodulation, [], [f428, f895])).
fof(f428, plain, ~ (op1(e14, e10) = op1(e14, e14)), inference(cnf_transformation, [], [f7])).
fof(f6064, plain, (~ spl68_28 | ~ spl68_29), inference(avatar_contradiction_clause, [], [f6063])).
fof(f6063, plain, ($false | (~ spl68_28 | ~ spl68_29)), inference(subsumption_resolution, [], [f6062, f542])).
fof(f542, plain, ~ (e12 = e13), inference(cnf_transformation, [], [f9])).
fof(f6062, plain, ((e12 = e13) | (~ spl68_28 | ~ spl68_29)), inference(forward_demodulation, [], [f1008, f1004])).
fof(f1008, plain, ((e13 = op1(e13, e14)) | ~ spl68_29), inference(avatar_component_clause, [], [f1006])).
fof(f1006, plain, (spl68_29 <=> (e13 = op1(e13, e14))), introduced(avatar_definition, [new_symbols(naming, [spl68_29])])).
fof(f6033, plain, (~ spl68_2 | ~ spl68_130), inference(avatar_contradiction_clause, [], [f6032])).
fof(f6032, plain, ($false | (~ spl68_2 | ~ spl68_130)), inference(subsumption_resolution, [], [f6031, f541])).
fof(f541, plain, ~ (e11 = e14), inference(cnf_transformation, [], [f9])).
fof(f6031, plain, ((e11 = e14) | (~ spl68_2 | ~ spl68_130)), inference(forward_demodulation, [], [f6009, f895])).
fof(f6009, plain, ((e14 = op1(e14, e14)) | ~ spl68_130), inference(backward_demodulation, [], [f197, f1432])).
fof(f1432, plain, ((e14 = unit1) | ~ spl68_130), inference(avatar_component_clause, [], [f1430])).
fof(f1430, plain, (spl68_130 <=> (e14 = unit1)), introduced(avatar_definition, [new_symbols(naming, [spl68_130])])).
fof(f5999, plain, (~ spl68_132 | ~ spl68_133), inference(avatar_contradiction_clause, [], [f5998])).
fof(f5998, plain, ($false | (~ spl68_132 | ~ spl68_133)), inference(subsumption_resolution, [], [f5997, f549])).
fof(f5997, plain, ((e21 = e22) | (~ spl68_132 | ~ spl68_133)), inference(forward_demodulation, [], [f1495, f1491])).
fof(f1495, plain, ((e22 = op2(e24, e24)) | ~ spl68_133), inference(avatar_component_clause, [], [f1493])).
fof(f1493, plain, (spl68_133 <=> (e22 = op2(e24, e24))), introduced(avatar_definition, [new_symbols(naming, [spl68_133])])).
fof(f5980, plain, (~ spl68_341 | ~ spl68_167), inference(avatar_split_clause, [], [f5976, f1636, f2665])).
fof(f5976, plain, (~ (e21 = h3(e11)) | ~ spl68_167), inference(backward_demodulation, [], [f2453, f1638])).
fof(f2453, plain, ~ (op2(e23, e22) = h3(e11)), inference(backward_demodulation, [], [f462, f750])).
fof(f462, plain, ~ (op2(e22, e22) = op2(e23, e22)), inference(cnf_transformation, [], [f8])).
fof(f5973, plain, (~ spl68_179 | ~ spl68_180), inference(avatar_contradiction_clause, [], [f5972])).
fof(f5972, plain, ($false | (~ spl68_179 | ~ spl68_180)), inference(subsumption_resolution, [], [f5971, f554])).
fof(f5971, plain, ((e23 = e24) | (~ spl68_179 | ~ spl68_180)), inference(backward_demodulation, [], [f1692, f1688])).
fof(f5951, plain, (~ spl68_206 | ~ spl68_208), inference(avatar_contradiction_clause, [], [f5950])).
fof(f5950, plain, ($false | (~ spl68_206 | ~ spl68_208)), inference(subsumption_resolution, [], [f5949, f546])).
fof(f546, plain, ~ (e20 = e22), inference(cnf_transformation, [], [f10])).
fof(f5949, plain, ((e20 = e22) | (~ spl68_206 | ~ spl68_208)), inference(forward_demodulation, [], [f1810, f1802])).
fof(f1810, plain, ((e22 = op2(e21, e24)) | ~ spl68_208), inference(avatar_component_clause, [], [f1808])).
fof(f1808, plain, (spl68_208 <=> (e22 = op2(e21, e24))), introduced(avatar_definition, [new_symbols(naming, [spl68_208])])).
fof(f5897, plain, (~ spl68_242 | ~ spl68_256), inference(avatar_contradiction_clause, [], [f5896])).
fof(f5896, plain, ($false | (~ spl68_242 | ~ spl68_256)), inference(subsumption_resolution, [], [f5895, f549])).
fof(f5895, plain, ((e21 = e22) | (~ spl68_242 | ~ spl68_256)), inference(forward_demodulation, [], [f5882, f1953])).
fof(f1953, plain, ((e21 = op2(e20, e22)) | ~ spl68_242), inference(avatar_component_clause, [], [f1951])).
fof(f1951, plain, (spl68_242 <=> (e21 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl68_242])])).
fof(f5882, plain, ((e22 = op2(e20, e22)) | ~ spl68_256), inference(backward_demodulation, [], [f278, f2012])).
fof(f5827, plain, (~ spl68_169 | ~ spl68_164), inference(avatar_split_clause, [], [f5822, f1623, f1644])).
fof(f1623, plain, (spl68_164 <=> (e23 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl68_164])])).
fof(f5822, plain, (~ (e23 = op2(e23, e22)) | ~ spl68_164), inference(backward_demodulation, [], [f2468, f5818])).
fof(f5818, plain, ((e23 = h4(e11)) | ~ spl68_164), inference(forward_demodulation, [], [f755, f1625])).
fof(f1625, plain, ((e23 = op2(e23, e23)) | ~ spl68_164), inference(avatar_component_clause, [], [f1623])).
fof(f2468, plain, ~ (op2(e23, e22) = h4(e11)), inference(backward_demodulation, [], [f522, f755])).
fof(f522, plain, ~ (op2(e23, e22) = op2(e23, e23)), inference(cnf_transformation, [], [f8])).
fof(f5826, plain, (~ spl68_239 | ~ spl68_164), inference(avatar_split_clause, [], [f5819, f1623, f1938])).
fof(f5819, plain, (~ (e23 = op2(e20, e23)) | ~ spl68_164), inference(backward_demodulation, [], [f2462, f5818])).
fof(f5799, plain, (~ spl68_137 | ~ spl68_314), inference(avatar_split_clause, [], [f5550, f2524, f1510])).
fof(f1510, plain, (spl68_137 <=> (e21 = op2(e24, e23))), introduced(avatar_definition, [new_symbols(naming, [spl68_137])])).
fof(f5550, plain, (~ (e21 = op2(e24, e23)) | ~ spl68_314), inference(forward_demodulation, [], [f2487, f2525])).
fof(f2487, plain, ~ (op2(e24, e23) = h5(e11)), inference(backward_demodulation, [], [f534, f760])).
fof(f534, plain, ~ (op2(e24, e23) = op2(e24, e24)), inference(cnf_transformation, [], [f8])).
fof(f5794, plain, (~ spl68_104 | ~ spl68_124), inference(avatar_split_clause, [], [f5793, f1405, f1321])).
fof(f5793, plain, (~ (e13 = op1(e10, e14)) | ~ spl68_124), inference(forward_demodulation, [], [f388, f1407])).
fof(f388, plain, ~ (op1(e10, e10) = op1(e10, e14)), inference(cnf_transformation, [], [f7])).
fof(f5764, plain, (~ spl68_6 | ~ spl68_8), inference(avatar_contradiction_clause, [], [f5763])).
fof(f5763, plain, ($false | (~ spl68_6 | ~ spl68_8)), inference(subsumption_resolution, [], [f5762, f536])).
fof(f536, plain, ~ (e10 = e12), inference(cnf_transformation, [], [f9])).
fof(f5762, plain, ((e10 = e12) | (~ spl68_6 | ~ spl68_8)), inference(backward_demodulation, [], [f920, f912])).
fof(f5741, plain, (~ spl68_93 | ~ spl68_95), inference(avatar_contradiction_clause, [], [f5740])).
fof(f5740, plain, ($false | (~ spl68_93 | ~ spl68_95)), inference(subsumption_resolution, [], [f5739, f543])).
fof(f5739, plain, ((e12 = e14) | (~ spl68_93 | ~ spl68_95)), inference(backward_demodulation, [], [f1285, f1277])).
fof(f5711, plain, (~ spl68_28 | ~ spl68_129), inference(avatar_contradiction_clause, [], [f5710])).
fof(f5710, plain, ($false | (~ spl68_28 | ~ spl68_129)), inference(subsumption_resolution, [], [f5709, f543])).
fof(f5709, plain, ((e12 = e14) | (~ spl68_28 | ~ spl68_129)), inference(forward_demodulation, [], [f5689, f1004])).
fof(f5689, plain, ((e14 = op1(e13, e14)) | ~ spl68_129), inference(backward_demodulation, [], [f196, f1428])).
fof(f1428, plain, ((e13 = unit1) | ~ spl68_129), inference(avatar_component_clause, [], [f1426])).
fof(f1426, plain, (spl68_129 <=> (e13 = unit1)), introduced(avatar_definition, [new_symbols(naming, [spl68_129])])).
fof(f5680, plain, (~ spl68_158 | ~ spl68_160), inference(avatar_contradiction_clause, [], [f5679])).
fof(f5679, plain, ($false | (~ spl68_158 | ~ spl68_160)), inference(subsumption_resolution, [], [f5678, f553])).
fof(f5678, plain, ((e22 = e24) | (~ spl68_158 | ~ spl68_160)), inference(forward_demodulation, [], [f1608, f1600])).
fof(f1608, plain, ((e24 = op2(e23, e24)) | ~ spl68_160), inference(avatar_component_clause, [], [f1606])).
fof(f1606, plain, (spl68_160 <=> (e24 = op2(e23, e24))), introduced(avatar_definition, [new_symbols(naming, [spl68_160])])).
fof(f5671, plain, (~ spl68_197 | ~ spl68_200), inference(avatar_contradiction_clause, [], [f5670])).
fof(f5670, plain, ($false | (~ spl68_197 | ~ spl68_200)), inference(subsumption_resolution, [], [f5669, f551])).
fof(f5669, plain, ((e21 = e24) | (~ spl68_197 | ~ spl68_200)), inference(backward_demodulation, [], [f1776, f1764])).
fof(f5659, plain, (spl68_356 | ~ spl68_223), inference(avatar_split_clause, [], [f5658, f1871, f2740])).
fof(f5658, plain, ((e22 = h2(e11)) | ~ spl68_223), inference(backward_demodulation, [], [f745, f1873])).
fof(f5657, plain, (~ spl68_227 | ~ spl68_229), inference(avatar_contradiction_clause, [], [f5656])).
fof(f5656, plain, ($false | (~ spl68_227 | ~ spl68_229)), inference(subsumption_resolution, [], [f5655, f550])).
fof(f550, plain, ~ (e21 = e23), inference(cnf_transformation, [], [f10])).
fof(f5655, plain, ((e21 = e23) | (~ spl68_227 | ~ spl68_229)), inference(forward_demodulation, [], [f1898, f1890])).
fof(f5646, plain, (~ spl68_184 | ~ spl68_258), inference(avatar_contradiction_clause, [], [f5645])).
fof(f5645, plain, ($false | (~ spl68_184 | ~ spl68_258)), inference(subsumption_resolution, [], [f5644, f554])).
fof(f5644, plain, ((e23 = e24) | (~ spl68_184 | ~ spl68_258)), inference(forward_demodulation, [], [f5624, f1709])).
fof(f5624, plain, ((e24 = op2(e22, e24)) | ~ spl68_258), inference(backward_demodulation, [], [f282, f2020])).
fof(f5640, plain, (~ spl68_191 | ~ spl68_258), inference(avatar_contradiction_clause, [], [f5639])).
fof(f5639, plain, ($false | (~ spl68_191 | ~ spl68_258)), inference(subsumption_resolution, [], [f5638, f546])).
fof(f5638, plain, ((e20 = e22) | (~ spl68_191 | ~ spl68_258)), inference(forward_demodulation, [], [f5621, f1739])).
fof(f1739, plain, ((e20 = op2(e22, e22)) | ~ spl68_191), inference(avatar_component_clause, [], [f1737])).
fof(f1737, plain, (spl68_191 <=> (e20 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl68_191])])).
fof(f5621, plain, ((e22 = op2(e22, e22)) | ~ spl68_258), inference(backward_demodulation, [], [f279, f2020])).
fof(f5613, plain, (spl68_329 | ~ spl68_171 | ~ spl68_326), inference(avatar_split_clause, [], [f5612, f2589, f1653, f2604])).
fof(f5612, plain, ((e20 = h4(e13)) | (~ spl68_171 | ~ spl68_326)), inference(forward_demodulation, [], [f5606, f1655])).
fof(f5606, plain, ((op2(e23, e21) = h4(e13)) | ~ spl68_326), inference(backward_demodulation, [], [f2470, f2590])).
fof(f2470, plain, (h4(e13) = op2(e23, h4(e11))), inference(backward_demodulation, [], [f757, f755])).
fof(f5537, plain, (~ spl68_123 | ~ spl68_113), inference(avatar_split_clause, [], [f5536, f1359, f1401])).
fof(f5536, plain, (~ (op1(e10, e10) = e12) | ~ spl68_113), inference(forward_demodulation, [], [f386, f1361])).
fof(f5528, plain, (~ spl68_119 | ~ spl68_19), inference(avatar_split_clause, [], [f5183, f964, f1384])).
fof(f1384, plain, (spl68_119 <=> (e13 = op1(e10, e11))), introduced(avatar_definition, [new_symbols(naming, [spl68_119])])).
fof(f5183, plain, (~ (e13 = op1(e10, e11)) | ~ spl68_19), inference(forward_demodulation, [], [f348, f966])).
fof(f348, plain, ~ (op1(e10, e11) = op1(e14, e11)), inference(cnf_transformation, [], [f7])).
fof(f5523, plain, (~ spl68_109 | ~ spl68_84), inference(avatar_split_clause, [], [f5522, f1237, f1342])).
fof(f1237, plain, (spl68_84 <=> (e13 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl68_84])])).
fof(f5522, plain, (~ (e13 = op1(e10, e13)) | ~ spl68_84), inference(forward_demodulation, [], [f365, f1239])).
fof(f1239, plain, ((e13 = op1(e11, e13)) | ~ spl68_84), inference(avatar_component_clause, [], [f1237])).
fof(f365, plain, ~ (op1(e10, e13) = op1(e11, e13)), inference(cnf_transformation, [], [f7])).
fof(f5497, plain, (~ spl68_59 | ~ spl68_84), inference(avatar_split_clause, [], [f5496, f1237, f1132])).
fof(f5496, plain, (~ (e13 = op1(e12, e13)) | ~ spl68_84), inference(forward_demodulation, [], [f369, f1239])).
fof(f369, plain, ~ (op1(e11, e13) = op1(e12, e13)), inference(cnf_transformation, [], [f7])).
fof(f5472, plain, (~ spl68_31 | ~ spl68_32), inference(avatar_contradiction_clause, [], [f5471])).
fof(f5471, plain, ($false | (~ spl68_31 | ~ spl68_32)), inference(subsumption_resolution, [], [f5470, f535])).
fof(f5470, plain, ((e10 = e11) | (~ spl68_31 | ~ spl68_32)), inference(backward_demodulation, [], [f1021, f1017])).
fof(f1017, plain, ((e10 = op1(e13, e13)) | ~ spl68_31), inference(avatar_component_clause, [], [f1015])).
fof(f5460, plain, (~ spl68_42 | ~ spl68_67), inference(avatar_split_clause, [], [f5459, f1166, f1061])).
fof(f1061, plain, (spl68_42 <=> (e11 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl68_42])])).
fof(f5459, plain, (~ (e11 = op1(e13, e11)) | ~ spl68_67), inference(backward_demodulation, [], [f352, f1168])).
fof(f352, plain, ~ (op1(e12, e11) = op1(e13, e11)), inference(cnf_transformation, [], [f7])).
fof(f5437, plain, (~ spl68_54 | ~ spl68_128), inference(avatar_contradiction_clause, [], [f5436])).
fof(f5436, plain, ($false | (~ spl68_54 | ~ spl68_128)), inference(subsumption_resolution, [], [f5435, f544])).
fof(f544, plain, ~ (e13 = e14), inference(cnf_transformation, [], [f9])).
fof(f5435, plain, ((e13 = e14) | (~ spl68_54 | ~ spl68_128)), inference(forward_demodulation, [], [f5427, f1113])).
fof(f5427, plain, ((e14 = op1(e12, e14)) | ~ spl68_128), inference(backward_demodulation, [], [f196, f1424])).
fof(f5433, plain, (~ spl68_57 | ~ spl68_128), inference(avatar_contradiction_clause, [], [f5432])).
fof(f5432, plain, ($false | (~ spl68_57 | ~ spl68_128)), inference(subsumption_resolution, [], [f5431, f540])).
fof(f5431, plain, ((e11 = e13) | (~ spl68_57 | ~ spl68_128)), inference(forward_demodulation, [], [f5425, f1126])).
fof(f1126, plain, ((e11 = op1(e12, e13)) | ~ spl68_57), inference(avatar_component_clause, [], [f1124])).
fof(f1124, plain, (spl68_57 <=> (e11 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl68_57])])).
fof(f5425, plain, ((e13 = op1(e12, e13)) | ~ spl68_128), inference(backward_demodulation, [], [f194, f1424])).
fof(f5416, plain, (~ spl68_144 | ~ spl68_149), inference(avatar_split_clause, [], [f5414, f1560, f1539])).
fof(f1539, plain, (spl68_144 <=> (e23 = op2(e24, e22))), introduced(avatar_definition, [new_symbols(naming, [spl68_144])])).
fof(f5414, plain, (~ (e23 = op2(e24, e22)) | ~ spl68_149), inference(backward_demodulation, [], [f529, f1562])).
fof(f529, plain, ~ (op2(e24, e21) = op2(e24, e22)), inference(cnf_transformation, [], [f8])).
fof(f5401, plain, (~ spl68_163 | spl68_321), inference(avatar_contradiction_clause, [], [f5400])).
fof(f5400, plain, ($false | (~ spl68_163 | spl68_321)), inference(subsumption_resolution, [], [f5399, f2566])).
fof(f2566, plain, (~ (e22 = h4(e11)) | spl68_321), inference(avatar_component_clause, [], [f2564])).
fof(f2564, plain, (spl68_321 <=> (e22 = h4(e11))), introduced(avatar_definition, [new_symbols(naming, [spl68_321])])).
fof(f5399, plain, ((e22 = h4(e11)) | ~ spl68_163), inference(backward_demodulation, [], [f755, f1621])).
fof(f1621, plain, ((e22 = op2(e23, e23)) | ~ spl68_163), inference(avatar_component_clause, [], [f1619])).
fof(f1619, plain, (spl68_163 <=> (e22 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl68_163])])).
fof(f5397, plain, (~ spl68_167 | ~ spl68_169), inference(avatar_contradiction_clause, [], [f5396])).
fof(f5396, plain, ($false | (~ spl68_167 | ~ spl68_169)), inference(subsumption_resolution, [], [f5395, f550])).
fof(f5395, plain, ((e21 = e23) | (~ spl68_167 | ~ spl68_169)), inference(backward_demodulation, [], [f1646, f1638])).
fof(f5393, plain, (~ spl68_171 | ~ spl68_175), inference(avatar_contradiction_clause, [], [f5392])).
fof(f5392, plain, ($false | (~ spl68_171 | ~ spl68_175)), inference(subsumption_resolution, [], [f5391, f548])).
fof(f5391, plain, ((e20 = e24) | (~ spl68_171 | ~ spl68_175)), inference(backward_demodulation, [], [f1671, f1655])).
fof(f1671, plain, ((e24 = op2(e23, e21)) | ~ spl68_175), inference(avatar_component_clause, [], [f1669])).
fof(f1669, plain, (spl68_175 <=> (e24 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl68_175])])).
fof(f5356, plain, (~ spl68_206 | ~ spl68_257), inference(avatar_contradiction_clause, [], [f5355])).
fof(f5355, plain, ($false | (~ spl68_206 | ~ spl68_257)), inference(subsumption_resolution, [], [f5354, f548])).
fof(f5354, plain, ((e20 = e24) | (~ spl68_206 | ~ spl68_257)), inference(forward_demodulation, [], [f5342, f1802])).
fof(f5342, plain, ((e24 = op2(e21, e24)) | ~ spl68_257), inference(backward_demodulation, [], [f282, f2016])).
fof(f2016, plain, ((e21 = unit2) | ~ spl68_257), inference(avatar_component_clause, [], [f2014])).
fof(f2014, plain, (spl68_257 <=> (e21 = unit2)), introduced(avatar_definition, [new_symbols(naming, [spl68_257])])).
fof(f5329, plain, (~ spl68_248 | ~ spl68_250), inference(avatar_contradiction_clause, [], [f5328])).
fof(f5328, plain, ($false | (~ spl68_248 | ~ spl68_250)), inference(subsumption_resolution, [], [f5327, f553])).
fof(f5327, plain, ((e22 = e24) | (~ spl68_248 | ~ spl68_250)), inference(backward_demodulation, [], [f1986, f1978])).
fof(f1986, plain, ((e24 = op2(e20, e21)) | ~ spl68_250), inference(avatar_component_clause, [], [f1984])).
fof(f5326, plain, (~ spl68_235 | ~ spl68_250), inference(avatar_split_clause, [], [f5325, f1984, f1921])).
fof(f5325, plain, (~ (e24 = op2(e20, e24)) | ~ spl68_250), inference(backward_demodulation, [], [f491, f1986])).
fof(f491, plain, ~ (op2(e20, e21) = op2(e20, e24)), inference(cnf_transformation, [], [f8])).
fof(f5281, plain, (~ spl68_222 | ~ spl68_227), inference(avatar_contradiction_clause, [], [f5280])).
fof(f5280, plain, ($false | (~ spl68_222 | ~ spl68_227)), inference(subsumption_resolution, [], [f4604, f5018])).
fof(f5018, plain, (~ (e21 = h2(e11)) | ~ spl68_227), inference(backward_demodulation, [], [f2444, f1890])).
fof(f2444, plain, ~ (op2(e21, e20) = h2(e11)), inference(backward_demodulation, [], [f495, f745])).
fof(f495, plain, ~ (op2(e21, e20) = op2(e21, e21)), inference(cnf_transformation, [], [f8])).
fof(f4604, plain, ((e21 = h2(e11)) | ~ spl68_222), inference(backward_demodulation, [], [f745, f1869])).
fof(f1869, plain, ((e21 = op2(e21, e21)) | ~ spl68_222), inference(avatar_component_clause, [], [f1867])).
fof(f1867, plain, (spl68_222 <=> (e21 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl68_222])])).
fof(f5265, plain, (~ spl68_183 | ~ spl68_310), inference(avatar_split_clause, [], [f5264, f2502, f1703])).
fof(f1703, plain, (spl68_183 <=> (e22 = op2(e22, e24))), introduced(avatar_definition, [new_symbols(naming, [spl68_183])])).
fof(f5264, plain, (~ (e22 = op2(e22, e24)) | ~ spl68_310), inference(forward_demodulation, [], [f2476, f2503])).
fof(f2476, plain, ~ (op2(e22, e24) = h5(e12)), inference(backward_demodulation, [], [f482, f2473])).
fof(f2473, plain, (op2(e23, e24) = h5(e12)), inference(forward_demodulation, [], [f761, f737])).
fof(f761, plain, (op2(op2(e24, op2(e24, e24)), e24) = h5(e12)), inference(cnf_transformation, [], [f20])).
fof(f482, plain, ~ (op2(e22, e24) = op2(e23, e24)), inference(cnf_transformation, [], [f8])).
fof(f5238, plain, (~ spl68_115 | ~ spl68_15), inference(avatar_split_clause, [], [f5237, f947, f1367])).
fof(f1367, plain, (spl68_115 <=> (e14 = op1(e10, e12))), introduced(avatar_definition, [new_symbols(naming, [spl68_115])])).
fof(f5237, plain, (~ (e14 = op1(e10, e12)) | ~ spl68_15), inference(forward_demodulation, [], [f358, f949])).
fof(f5203, plain, (~ spl68_43 | ~ spl68_28), inference(avatar_split_clause, [], [f4789, f1002, f1065])).
fof(f1065, plain, (spl68_43 <=> (e12 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl68_43])])).
fof(f4789, plain, (~ (e12 = op1(e13, e11)) | ~ spl68_28), inference(forward_demodulation, [], [f421, f1004])).
fof(f421, plain, ~ (op1(e13, e11) = op1(e13, e14)), inference(cnf_transformation, [], [f7])).
fof(f5197, plain, (~ spl68_19 | spl68_512), inference(avatar_contradiction_clause, [], [f5196])).
fof(f5196, plain, ($false | (~ spl68_19 | spl68_512)), inference(trivial_inequality_removal, [], [f5195])).
fof(f5195, plain, (~ (h1(e13) = h1(e13)) | (~ spl68_19 | spl68_512)), inference(forward_demodulation, [], [f3702, f966])).
fof(f3702, plain, (~ (h1(e13) = h1(op1(e14, e11))) | spl68_512), inference(avatar_component_clause, [], [f3700])).
fof(f3700, plain, (spl68_512 <=> (h1(e13) = h1(op1(e14, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl68_512])])).
fof(f5187, plain, (~ spl68_69 | ~ spl68_19), inference(avatar_split_clause, [], [f5186, f964, f1174])).
fof(f1174, plain, (spl68_69 <=> (e13 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl68_69])])).
fof(f5186, plain, (~ (e13 = op1(e12, e11)) | ~ spl68_19), inference(forward_demodulation, [], [f353, f966])).
fof(f353, plain, ~ (op1(e12, e11) = op1(e14, e11)), inference(cnf_transformation, [], [f7])).
fof(f5180, plain, (~ spl68_7 | ~ spl68_2), inference(avatar_split_clause, [], [f5179, f893, f914])).
fof(f914, plain, (spl68_7 <=> (e11 = op1(e14, e13))), introduced(avatar_definition, [new_symbols(naming, [spl68_7])])).
fof(f5179, plain, (~ (e11 = op1(e14, e13)) | ~ spl68_2), inference(forward_demodulation, [], [f434, f895])).
fof(f434, plain, ~ (op1(e14, e13) = op1(e14, e14)), inference(cnf_transformation, [], [f7])).
fof(f5169, plain, (~ spl68_18 | ~ spl68_19), inference(avatar_contradiction_clause, [], [f5168])).
fof(f5168, plain, ($false | (~ spl68_18 | ~ spl68_19)), inference(subsumption_resolution, [], [f5167, f542])).
fof(f5167, plain, ((e12 = e13) | (~ spl68_18 | ~ spl68_19)), inference(forward_demodulation, [], [f966, f962])).
fof(f962, plain, ((e12 = op1(e14, e11)) | ~ spl68_18), inference(avatar_component_clause, [], [f960])).
fof(f960, plain, (spl68_18 <=> (e12 = op1(e14, e11))), introduced(avatar_definition, [new_symbols(naming, [spl68_18])])).
fof(f5112, plain, (~ spl68_76 | ~ spl68_127), inference(avatar_contradiction_clause, [], [f5111])).
fof(f5111, plain, ($false | (~ spl68_76 | ~ spl68_127)), inference(subsumption_resolution, [], [f5110, f538])).
fof(f5110, plain, ((e10 = e14) | (~ spl68_76 | ~ spl68_127)), inference(forward_demodulation, [], [f5093, f1206])).
fof(f5093, plain, ((e14 = op1(e11, e14)) | ~ spl68_127), inference(backward_demodulation, [], [f196, f1420])).
fof(f1420, plain, ((e11 = unit1) | ~ spl68_127), inference(avatar_component_clause, [], [f1418])).
fof(f1418, plain, (spl68_127 <=> (e11 = unit1)), introduced(avatar_definition, [new_symbols(naming, [spl68_127])])).
fof(f5071, plain, (~ spl68_149 | ~ spl68_154), inference(avatar_split_clause, [], [f5068, f1581, f1560])).
fof(f1581, plain, (spl68_154 <=> (e23 = op2(e24, e20))), introduced(avatar_definition, [new_symbols(naming, [spl68_154])])).
fof(f5068, plain, (~ (e23 = op2(e24, e21)) | ~ spl68_154), inference(backward_demodulation, [], [f525, f1583])).
fof(f1583, plain, ((e23 = op2(e24, e20)) | ~ spl68_154), inference(avatar_component_clause, [], [f1581])).
fof(f525, plain, ~ (op2(e24, e20) = op2(e24, e21)), inference(cnf_transformation, [], [f8])).
fof(f5063, plain, (spl68_326 | ~ spl68_162), inference(avatar_split_clause, [], [f5062, f1615, f2589])).
fof(f5062, plain, ((e21 = h4(e11)) | ~ spl68_162), inference(backward_demodulation, [], [f755, f1617])).
fof(f5058, plain, (~ spl68_170 | ~ spl68_175), inference(avatar_split_clause, [], [f5055, f1669, f1648])).
fof(f5055, plain, (~ (e24 = op2(e23, e22)) | ~ spl68_175), inference(backward_demodulation, [], [f519, f1671])).
fof(f519, plain, ~ (op2(e23, e21) = op2(e23, e22)), inference(cnf_transformation, [], [f8])).
fof(f5028, plain, (~ spl68_206 | ~ spl68_209), inference(avatar_contradiction_clause, [], [f5027])).
fof(f5027, plain, ($false | (~ spl68_206 | ~ spl68_209)), inference(subsumption_resolution, [], [f5025, f547])).
fof(f5025, plain, ((e20 = e23) | (~ spl68_206 | ~ spl68_209)), inference(backward_demodulation, [], [f1814, f1802])).
fof(f1814, plain, ((e23 = op2(e21, e24)) | ~ spl68_209), inference(avatar_component_clause, [], [f1812])).
fof(f1812, plain, (spl68_209 <=> (e23 = op2(e21, e24))), introduced(avatar_definition, [new_symbols(naming, [spl68_209])])).
fof(f5023, plain, (~ spl68_217 | ~ spl68_218), inference(avatar_contradiction_clause, [], [f5022])).
fof(f5022, plain, ($false | (~ spl68_217 | ~ spl68_218)), inference(subsumption_resolution, [], [f5021, f549])).
fof(f5021, plain, ((e21 = e22) | (~ spl68_217 | ~ spl68_218)), inference(forward_demodulation, [], [f1852, f1848])).
fof(f1852, plain, ((e22 = op2(e21, e22)) | ~ spl68_218), inference(avatar_component_clause, [], [f1850])).
fof(f1850, plain, (spl68_218 <=> (e22 = op2(e21, e22))), introduced(avatar_definition, [new_symbols(naming, [spl68_218])])).
fof(f5014, plain, (~ spl68_251 | ~ spl68_253), inference(avatar_contradiction_clause, [], [f5013])).
fof(f5013, plain, ($false | (~ spl68_251 | ~ spl68_253)), inference(subsumption_resolution, [], [f5012, f546])).
fof(f5012, plain, ((e20 = e22) | (~ spl68_251 | ~ spl68_253)), inference(backward_demodulation, [], [f1999, f1991])).
fof(f1999, plain, ((op2(e20, e20) = e22) | ~ spl68_253), inference(avatar_component_clause, [], [f1997])).
fof(f1997, plain, (spl68_253 <=> (op2(e20, e20) = e22)), introduced(avatar_definition, [new_symbols(naming, [spl68_253])])).
fof(f5008, plain, (~ spl68_132 | ~ spl68_260), inference(avatar_contradiction_clause, [], [f5007])).
fof(f5007, plain, ($false | (~ spl68_132 | ~ spl68_260)), inference(subsumption_resolution, [], [f5006, f551])).
fof(f5006, plain, ((e21 = e24) | (~ spl68_132 | ~ spl68_260)), inference(forward_demodulation, [], [f4985, f1491])).
fof(f4985, plain, ((e24 = op2(e24, e24)) | ~ spl68_260), inference(backward_demodulation, [], [f283, f2028])).
fof(f2028, plain, ((e24 = unit2) | ~ spl68_260), inference(avatar_component_clause, [], [f2026])).
fof(f2026, plain, (spl68_260 <=> (e24 = unit2)), introduced(avatar_definition, [new_symbols(naming, [spl68_260])])).
fof(f4975, plain, (~ spl68_168 | ~ spl68_310), inference(avatar_split_clause, [], [f4971, f2502, f1640])).
fof(f1640, plain, (spl68_168 <=> (e22 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl68_168])])).
fof(f4971, plain, (~ (e22 = op2(e23, e22)) | ~ spl68_310), inference(backward_demodulation, [], [f2480, f2503])).
fof(f2480, plain, ~ (op2(e23, e22) = h5(e12)), inference(backward_demodulation, [], [f523, f2473])).
fof(f523, plain, ~ (op2(e23, e22) = op2(e23, e24)), inference(cnf_transformation, [], [f8])).
fof(f4974, plain, (~ spl68_173 | ~ spl68_310), inference(avatar_split_clause, [], [f4970, f2502, f1661])).
fof(f1661, plain, (spl68_173 <=> (e22 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl68_173])])).
fof(f4970, plain, (~ (e22 = op2(e23, e21)) | ~ spl68_310), inference(backward_demodulation, [], [f2479, f2503])).
fof(f2479, plain, ~ (op2(e23, e21) = h5(e12)), inference(backward_demodulation, [], [f521, f2473])).
fof(f521, plain, ~ (op2(e23, e21) = op2(e23, e24)), inference(cnf_transformation, [], [f8])).
fof(f4973, plain, (~ spl68_178 | ~ spl68_310), inference(avatar_split_clause, [], [f4969, f2502, f1682])).
fof(f1682, plain, (spl68_178 <=> (e22 = op2(e23, e20))), introduced(avatar_definition, [new_symbols(naming, [spl68_178])])).
fof(f4969, plain, (~ (e22 = op2(e23, e20)) | ~ spl68_310), inference(backward_demodulation, [], [f2478, f2503])).
fof(f2478, plain, ~ (op2(e23, e20) = h5(e12)), inference(backward_demodulation, [], [f518, f2473])).
fof(f518, plain, ~ (op2(e23, e20) = op2(e23, e24)), inference(cnf_transformation, [], [f8])).
fof(f4965, plain, (spl68_149 | ~ spl68_314), inference(avatar_split_clause, [], [f4958, f2524, f1560])).
fof(f4958, plain, ((e23 = op2(e24, e21)) | ~ spl68_314), inference(backward_demodulation, [], [f2489, f2525])).
fof(f2489, plain, (e23 = op2(e24, h5(e11))), inference(backward_demodulation, [], [f737, f760])).
fof(f4964, plain, (spl68_206 | ~ spl68_314), inference(avatar_split_clause, [], [f4957, f2524, f1800])).
fof(f4957, plain, ((e20 = op2(e21, e24)) | ~ spl68_314), inference(backward_demodulation, [], [f2488, f2525])).
fof(f4945, plain, (~ spl68_153 | ~ spl68_371), inference(avatar_split_clause, [], [f4938, f2816, f1577])).
fof(f2816, plain, (spl68_371 <=> (e22 = h1(e11))), introduced(avatar_definition, [new_symbols(naming, [spl68_371])])).
fof(f4938, plain, (~ (e22 = op2(e24, e20)) | ~ spl68_371), inference(backward_demodulation, [], [f2432, f2817])).
fof(f2817, plain, ((e22 = h1(e11)) | ~ spl68_371), inference(avatar_component_clause, [], [f2816])).
fof(f2432, plain, ~ (op2(e24, e20) = h1(e11)), inference(backward_demodulation, [], [f438, f740])).
fof(f438, plain, ~ (op2(e20, e20) = op2(e24, e20)), inference(cnf_transformation, [], [f8])).
fof(f4928, plain, (~ spl68_172 | ~ spl68_247), inference(avatar_split_clause, [], [f4927, f1972, f1657])).
fof(f4927, plain, (~ (e21 = op2(e23, e21)) | ~ spl68_247), inference(forward_demodulation, [], [f447, f1974])).
fof(f447, plain, ~ (op2(e20, e21) = op2(e23, e21)), inference(cnf_transformation, [], [f8])).
fof(f4903, plain, (~ spl68_227 | ~ spl68_217), inference(avatar_split_clause, [], [f4902, f1846, f1888])).
fof(f4902, plain, (~ (e21 = op2(e21, e20)) | ~ spl68_217), inference(forward_demodulation, [], [f496, f1848])).
fof(f496, plain, ~ (op2(e21, e20) = op2(e21, e22)), inference(cnf_transformation, [], [f8])).
fof(f4874, plain, (~ spl68_206 | ~ spl68_181), inference(avatar_split_clause, [], [f4873, f1695, f1800])).
fof(f1695, plain, (spl68_181 <=> (e20 = op2(e22, e24))), introduced(avatar_definition, [new_symbols(naming, [spl68_181])])).
fof(f4873, plain, (~ (e20 = op2(e21, e24)) | ~ spl68_181), inference(forward_demodulation, [], [f479, f1697])).
fof(f1697, plain, ((e20 = op2(e22, e24)) | ~ spl68_181), inference(avatar_component_clause, [], [f1695])).
fof(f479, plain, ~ (op2(e21, e24) = op2(e22, e24)), inference(cnf_transformation, [], [f8])).
fof(f4865, plain, (~ spl68_197 | ~ spl68_222), inference(avatar_contradiction_clause, [], [f4864])).
fof(f4864, plain, ($false | (~ spl68_197 | ~ spl68_222)), inference(subsumption_resolution, [], [f4604, f4632])).
fof(f4632, plain, (~ (e21 = h2(e11)) | ~ spl68_197), inference(backward_demodulation, [], [f2441, f1764])).
fof(f2441, plain, ~ (op2(e22, e21) = h2(e11)), inference(backward_demodulation, [], [f449, f745])).
fof(f449, plain, ~ (op2(e21, e21) = op2(e22, e21)), inference(cnf_transformation, [], [f8])).
fof(f4831, plain, (~ spl68_82 | ~ spl68_32), inference(avatar_split_clause, [], [f4830, f1019, f1229])).
fof(f1229, plain, (spl68_82 <=> (e11 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl68_82])])).
fof(f4830, plain, (~ (e11 = op1(e11, e13)) | ~ spl68_32), inference(forward_demodulation, [], [f370, f1021])).
fof(f4824, plain, (~ spl68_81 | ~ spl68_76), inference(avatar_split_clause, [], [f4823, f1204, f1225])).
fof(f1225, plain, (spl68_81 <=> (e10 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl68_81])])).
fof(f4823, plain, (~ (e10 = op1(e11, e13)) | ~ spl68_76), inference(forward_demodulation, [], [f404, f1206])).
fof(f404, plain, ~ (op1(e11, e13) = op1(e11, e14)), inference(cnf_transformation, [], [f7])).
fof(f4801, plain, (~ spl68_48 | ~ spl68_28), inference(avatar_split_clause, [], [f4800, f1002, f1086])).
fof(f1086, plain, (spl68_48 <=> (e12 = op1(e13, e10))), introduced(avatar_definition, [new_symbols(naming, [spl68_48])])).
fof(f4800, plain, (~ (e12 = op1(e13, e10)) | ~ spl68_28), inference(forward_demodulation, [], [f418, f1004])).
fof(f418, plain, ~ (op1(e13, e10) = op1(e13, e14)), inference(cnf_transformation, [], [f7])).
fof(f4799, plain, (~ spl68_46 | ~ spl68_71), inference(avatar_split_clause, [], [f4798, f1183, f1078])).
fof(f1078, plain, (spl68_46 <=> (e10 = op1(e13, e10))), introduced(avatar_definition, [new_symbols(naming, [spl68_46])])).
fof(f4798, plain, (~ (e10 = op1(e13, e10)) | ~ spl68_71), inference(forward_demodulation, [], [f342, f1185])).
fof(f342, plain, ~ (op1(e12, e10) = op1(e13, e10)), inference(cnf_transformation, [], [f7])).
fof(f4786, plain, (~ spl68_38 | ~ spl68_28), inference(avatar_split_clause, [], [f4785, f1002, f1044])).
fof(f1044, plain, (spl68_38 <=> (e12 = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl68_38])])).
fof(f4785, plain, (~ (e12 = op1(e13, e12)) | ~ spl68_28), inference(forward_demodulation, [], [f423, f1004])).
fof(f423, plain, ~ (op1(e13, e12) = op1(e13, e14)), inference(cnf_transformation, [], [f7])).
fof(f4784, plain, (~ spl68_37 | ~ spl68_87), inference(avatar_split_clause, [], [f4783, f1250, f1040])).
fof(f4783, plain, (~ (e11 = op1(e13, e12)) | ~ spl68_87), inference(forward_demodulation, [], [f360, f1252])).
fof(f4741, plain, (~ spl68_76 | ~ spl68_79), inference(avatar_contradiction_clause, [], [f4740])).
fof(f4740, plain, ($false | (~ spl68_76 | ~ spl68_79)), inference(subsumption_resolution, [], [f4739, f537])).
fof(f4739, plain, ((e10 = e13) | (~ spl68_76 | ~ spl68_79)), inference(backward_demodulation, [], [f1218, f1206])).
fof(f1218, plain, ((e13 = op1(e11, e14)) | ~ spl68_79), inference(avatar_component_clause, [], [f1216])).
fof(f1216, plain, (spl68_79 <=> (e13 = op1(e11, e14))), introduced(avatar_definition, [new_symbols(naming, [spl68_79])])).
fof(f4732, plain, (~ spl68_82 | ~ spl68_87), inference(avatar_split_clause, [], [f4730, f1250, f1229])).
fof(f4730, plain, (~ (e11 = op1(e11, e13)) | ~ spl68_87), inference(backward_demodulation, [], [f402, f1252])).
fof(f402, plain, ~ (op1(e11, e12) = op1(e11, e13)), inference(cnf_transformation, [], [f7])).
fof(f4708, plain, (~ spl68_13 | ~ spl68_128), inference(avatar_contradiction_clause, [], [f4707])).
fof(f4707, plain, ($false | (~ spl68_13 | ~ spl68_128)), inference(subsumption_resolution, [], [f4706, f543])).
fof(f4706, plain, ((e12 = e14) | (~ spl68_13 | ~ spl68_128)), inference(forward_demodulation, [], [f4693, f941])).
fof(f941, plain, ((e12 = op1(e14, e12)) | ~ spl68_13), inference(avatar_component_clause, [], [f939])).
fof(f4693, plain, ((e14 = op1(e14, e12)) | ~ spl68_128), inference(backward_demodulation, [], [f197, f1424])).
fof(f4697, plain, (~ spl68_66 | ~ spl68_128), inference(avatar_contradiction_clause, [], [f4696])).
fof(f4696, plain, ($false | (~ spl68_66 | ~ spl68_128)), inference(subsumption_resolution, [], [f4695, f535])).
fof(f4695, plain, ((e10 = e11) | (~ spl68_66 | ~ spl68_128)), inference(forward_demodulation, [], [f4686, f1164])).
fof(f1164, plain, ((e10 = op1(e12, e11)) | ~ spl68_66), inference(avatar_component_clause, [], [f1162])).
fof(f1162, plain, (spl68_66 <=> (e10 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl68_66])])).
fof(f4686, plain, ((e11 = op1(e12, e11)) | ~ spl68_128), inference(backward_demodulation, [], [f190, f1424])).
fof(f4672, plain, (spl68_331 | ~ spl68_161), inference(avatar_split_clause, [], [f4671, f1611, f2614])).
fof(f1611, plain, (spl68_161 <=> (e20 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl68_161])])).
fof(f4671, plain, ((e20 = h4(e11)) | ~ spl68_161), inference(backward_demodulation, [], [f755, f1613])).
fof(f1613, plain, ((e20 = op2(e23, e23)) | ~ spl68_161), inference(avatar_component_clause, [], [f1611])).
fof(f4656, plain, (~ spl68_174 | ~ spl68_179), inference(avatar_split_clause, [], [f4650, f1686, f1665])).
fof(f1665, plain, (spl68_174 <=> (e23 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl68_174])])).
fof(f4650, plain, (~ (e23 = op2(e23, e21)) | ~ spl68_179), inference(backward_demodulation, [], [f515, f1688])).
fof(f515, plain, ~ (op2(e23, e20) = op2(e23, e21)), inference(cnf_transformation, [], [f8])).
fof(f4635, plain, (~ spl68_187 | ~ spl68_197), inference(avatar_split_clause, [], [f4631, f1762, f1720])).
fof(f1720, plain, (spl68_187 <=> (e21 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl68_187])])).
fof(f4631, plain, (~ (e21 = op2(e22, e23)) | ~ spl68_197), inference(backward_demodulation, [], [f510, f1764])).
fof(f510, plain, ~ (op2(e22, e21) = op2(e22, e23)), inference(cnf_transformation, [], [f8])).
fof(f4619, plain, (~ spl68_432 | ~ spl68_215), inference(avatar_split_clause, [], [f4616, f1837, f3224])).
fof(f4616, plain, (~ (e24 = h4(e11)) | ~ spl68_215), inference(backward_demodulation, [], [f2463, f1839])).
fof(f2463, plain, ~ (op2(e21, e23) = h4(e11)), inference(backward_demodulation, [], [f470, f755])).
fof(f470, plain, ~ (op2(e21, e23) = op2(e23, e23)), inference(cnf_transformation, [], [f8])).
fof(f4613, plain, (~ spl68_212 | ~ spl68_217), inference(avatar_split_clause, [], [f4609, f1846, f1825])).
fof(f4609, plain, (~ (e21 = op2(e21, e23)) | ~ spl68_217), inference(backward_demodulation, [], [f502, f1848])).
fof(f502, plain, ~ (op2(e21, e22) = op2(e21, e23)), inference(cnf_transformation, [], [f8])).
fof(f4601, plain, (~ spl68_214 | ~ spl68_239), inference(avatar_split_clause, [], [f4596, f1938, f1833])).
fof(f1833, plain, (spl68_214 <=> (e23 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl68_214])])).
fof(f4596, plain, (~ (e23 = op2(e21, e23)) | ~ spl68_239), inference(backward_demodulation, [], [f465, f1940])).
fof(f465, plain, ~ (op2(e20, e23) = op2(e21, e23)), inference(cnf_transformation, [], [f8])).
fof(f4599, plain, (~ spl68_337 | ~ spl68_239 | ~ spl68_346), inference(avatar_split_clause, [], [f4594, f2690, f1938, f2644])).
fof(f2690, plain, (spl68_346 <=> (e20 = h3(e11))), introduced(avatar_definition, [new_symbols(naming, [spl68_346])])).
fof(f4594, plain, (~ (e23 = h3(e10)) | (~ spl68_239 | ~ spl68_346)), inference(backward_demodulation, [], [f4554, f1940])).
fof(f4554, plain, (~ (op2(e20, e23) = h3(e10)) | ~ spl68_346), inference(backward_demodulation, [], [f492, f4541])).
fof(f4541, plain, ((op2(e20, e22) = h3(e10)) | ~ spl68_346), inference(backward_demodulation, [], [f2460, f2691])).
fof(f2691, plain, ((e20 = h3(e11)) | ~ spl68_346), inference(avatar_component_clause, [], [f2690])).
fof(f492, plain, ~ (op2(e20, e22) = op2(e20, e23)), inference(cnf_transformation, [], [f8])).
fof(f4580, plain, (spl68_371 | ~ spl68_253), inference(avatar_split_clause, [], [f4579, f1997, f2816])).
fof(f4579, plain, ((e22 = h1(e11)) | ~ spl68_253), inference(backward_demodulation, [], [f740, f1999])).
fof(f4575, plain, (~ spl68_158 | ~ spl68_259), inference(avatar_contradiction_clause, [], [f4574])).
fof(f4574, plain, ($false | (~ spl68_158 | ~ spl68_259)), inference(subsumption_resolution, [], [f4573, f553])).
fof(f4573, plain, ((e22 = e24) | (~ spl68_158 | ~ spl68_259)), inference(forward_demodulation, [], [f4563, f1600])).
fof(f4563, plain, ((e24 = op2(e23, e24)) | ~ spl68_259), inference(backward_demodulation, [], [f282, f2024])).
fof(f2024, plain, ((e23 = unit2) | ~ spl68_259), inference(avatar_component_clause, [], [f2022])).
fof(f2022, plain, (spl68_259 <=> (e23 = unit2)), introduced(avatar_definition, [new_symbols(naming, [spl68_259])])).
fof(f4481, plain, (~ spl68_314 | ~ spl68_182), inference(avatar_split_clause, [], [f4480, f1699, f2524])).
fof(f1699, plain, (spl68_182 <=> (e21 = op2(e22, e24))), introduced(avatar_definition, [new_symbols(naming, [spl68_182])])).
fof(f4480, plain, (~ (e21 = h5(e11)) | ~ spl68_182), inference(forward_demodulation, [], [f2483, f1701])).
fof(f1701, plain, ((e21 = op2(e22, e24)) | ~ spl68_182), inference(avatar_component_clause, [], [f1699])).
fof(f2483, plain, ~ (op2(e22, e24) = h5(e11)), inference(backward_demodulation, [], [f483, f760])).
fof(f483, plain, ~ (op2(e22, e24) = op2(e24, e24)), inference(cnf_transformation, [], [f8])).
fof(f4470, plain, (~ spl68_321 | ~ spl68_158), inference(avatar_split_clause, [], [f4469, f1598, f2564])).
fof(f4469, plain, (~ (e22 = h4(e11)) | ~ spl68_158), inference(forward_demodulation, [], [f2469, f1600])).
fof(f2469, plain, ~ (op2(e23, e24) = h4(e11)), inference(backward_demodulation, [], [f524, f755])).
fof(f524, plain, ~ (op2(e23, e23) = op2(e23, e24)), inference(cnf_transformation, [], [f8])).
fof(f4468, plain, (spl68_310 | ~ spl68_158), inference(avatar_split_clause, [], [f4467, f1598, f2502])).
fof(f4467, plain, ((e22 = h5(e12)) | ~ spl68_158), inference(forward_demodulation, [], [f2473, f1600])).
fof(f4466, plain, (~ spl68_153 | ~ spl68_155), inference(avatar_contradiction_clause, [], [f4465])).
fof(f4465, plain, ($false | (~ spl68_153 | ~ spl68_155)), inference(subsumption_resolution, [], [f4464, f553])).
fof(f4464, plain, ((e22 = e24) | (~ spl68_153 | ~ spl68_155)), inference(forward_demodulation, [], [f1587, f1579])).
fof(f4446, plain, (~ spl68_314 | ~ spl68_142), inference(avatar_split_clause, [], [f4445, f1531, f2524])).
fof(f1531, plain, (spl68_142 <=> (e21 = op2(e24, e22))), introduced(avatar_definition, [new_symbols(naming, [spl68_142])])).
fof(f4445, plain, (~ (e21 = h5(e11)) | ~ spl68_142), inference(forward_demodulation, [], [f2486, f1533])).
fof(f1533, plain, ((e21 = op2(e24, e22)) | ~ spl68_142), inference(avatar_component_clause, [], [f1531])).
fof(f2486, plain, ~ (op2(e24, e22) = h5(e11)), inference(backward_demodulation, [], [f533, f760])).
fof(f533, plain, ~ (op2(e24, e22) = op2(e24, e24)), inference(cnf_transformation, [], [f8])).
fof(f4438, plain, (spl68_314 | ~ spl68_132), inference(avatar_split_clause, [], [f4437, f1489, f2524])).
fof(f4437, plain, ((e21 = h5(e11)) | ~ spl68_132), inference(forward_demodulation, [], [f760, f1491])).
fof(f4432, plain, (~ spl68_117 | ~ spl68_42), inference(avatar_split_clause, [], [f4431, f1061, f1376])).
fof(f4431, plain, (~ (e11 = op1(e10, e11)) | ~ spl68_42), inference(forward_demodulation, [], [f347, f1063])).
fof(f1063, plain, ((e11 = op1(e13, e11)) | ~ spl68_42), inference(avatar_component_clause, [], [f1061])).
fof(f347, plain, ~ (op1(e10, e11) = op1(e13, e11)), inference(cnf_transformation, [], [f7])).
fof(f4430, plain, (~ spl68_116 | ~ spl68_111), inference(avatar_split_clause, [], [f4429, f1351, f1372])).
fof(f4429, plain, (~ (e10 = op1(e10, e11)) | ~ spl68_111), inference(forward_demodulation, [], [f389, f1353])).
fof(f4380, plain, (~ spl68_53 | ~ spl68_28), inference(avatar_split_clause, [], [f4379, f1002, f1107])).
fof(f1107, plain, (spl68_53 <=> (e12 = op1(e12, e14))), introduced(avatar_definition, [new_symbols(naming, [spl68_53])])).
fof(f4379, plain, (~ (e12 = op1(e12, e14)) | ~ spl68_28), inference(forward_demodulation, [], [f382, f1004])).
fof(f382, plain, ~ (op1(e12, e14) = op1(e13, e14)), inference(cnf_transformation, [], [f7])).
fof(f4378, plain, (~ spl68_52 | ~ spl68_2), inference(avatar_split_clause, [], [f4377, f893, f1103])).
fof(f1103, plain, (spl68_52 <=> (e11 = op1(e12, e14))), introduced(avatar_definition, [new_symbols(naming, [spl68_52])])).
fof(f4377, plain, (~ (e11 = op1(e12, e14)) | ~ spl68_2), inference(forward_demodulation, [], [f383, f895])).
fof(f383, plain, ~ (op1(e12, e14) = op1(e14, e14)), inference(cnf_transformation, [], [f7])).
fof(f4371, plain, (~ spl68_33 | ~ spl68_28), inference(avatar_split_clause, [], [f4370, f1002, f1023])).
fof(f1023, plain, (spl68_33 <=> (e12 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl68_33])])).
fof(f4370, plain, (~ (e12 = op1(e13, e13)) | ~ spl68_28), inference(forward_demodulation, [], [f424, f1004])).
fof(f424, plain, ~ (op1(e13, e13) = op1(e13, e14)), inference(cnf_transformation, [], [f7])).
fof(f4369, plain, (~ spl68_23 | ~ spl68_25), inference(avatar_contradiction_clause, [], [f4368])).
fof(f4368, plain, ($false | (~ spl68_23 | ~ spl68_25)), inference(subsumption_resolution, [], [f4367, f543])).
fof(f4367, plain, ((e12 = e14) | (~ spl68_23 | ~ spl68_25)), inference(forward_demodulation, [], [f991, f983])).
fof(f4360, plain, (spl68_19 | ~ spl68_2), inference(avatar_split_clause, [], [f4351, f893, f964])).
fof(f4351, plain, ((e13 = op1(e14, e11)) | ~ spl68_2), inference(backward_demodulation, [], [f733, f895])).
fof(f733, plain, (e13 = op1(e14, op1(e14, e14))), inference(cnf_transformation, [], [f14])).
fof(f14, plain, ((e13 = op1(e14, op1(e14, e14))) & (e12 = op1(op1(e14, op1(e14, e14)), e14)) & (e11 = op1(e14, e14)) & (e10 = op1(op1(e14, e14), e14))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG052+1.p', ax14)).
fof(f4357, plain, (~ spl68_12 | ~ spl68_2), inference(avatar_split_clause, [], [f4356, f893, f935])).
fof(f935, plain, (spl68_12 <=> (e11 = op1(e14, e12))), introduced(avatar_definition, [new_symbols(naming, [spl68_12])])).
fof(f4356, plain, (~ (e11 = op1(e14, e12)) | ~ spl68_2), inference(forward_demodulation, [], [f433, f895])).
fof(f433, plain, ~ (op1(e14, e12) = op1(e14, e14)), inference(cnf_transformation, [], [f7])).
fof(f4355, plain, (~ spl68_2 | ~ spl68_16), inference(avatar_contradiction_clause, [], [f4354])).
fof(f4354, plain, ($false | (~ spl68_2 | ~ spl68_16)), inference(subsumption_resolution, [], [f4353, f537])).
fof(f4353, plain, ((e10 = e13) | (~ spl68_2 | ~ spl68_16)), inference(forward_demodulation, [], [f4351, f954])).
fof(f954, plain, ((e10 = op1(e14, e11)) | ~ spl68_16), inference(avatar_component_clause, [], [f952])).
fof(f952, plain, (spl68_16 <=> (e10 = op1(e14, e11))), introduced(avatar_definition, [new_symbols(naming, [spl68_16])])).
fof(f4352, plain, (spl68_76 | ~ spl68_2), inference(avatar_split_clause, [], [f4350, f893, f1204])).
fof(f4350, plain, ((e10 = op1(e11, e14)) | ~ spl68_2), inference(backward_demodulation, [], [f730, f895])).
fof(f730, plain, (e10 = op1(op1(e14, e14), e14)), inference(cnf_transformation, [], [f14])).
fof(f4321, plain, (~ spl68_5 | ~ spl68_25), inference(avatar_split_clause, [], [f4317, f989, f905])).
fof(f905, plain, (spl68_5 <=> (e14 = op1(e14, e14))), introduced(avatar_definition, [new_symbols(naming, [spl68_5])])).
fof(f4317, plain, (~ (e14 = op1(e14, e14)) | ~ spl68_25), inference(backward_demodulation, [], [f428, f991])).
fof(f4313, plain, (~ spl68_3 | ~ spl68_28), inference(avatar_split_clause, [], [f4312, f1002, f897])).
fof(f897, plain, (spl68_3 <=> (e12 = op1(e14, e14))), introduced(avatar_definition, [new_symbols(naming, [spl68_3])])).
fof(f4312, plain, (~ (e12 = op1(e14, e14)) | ~ spl68_28), inference(backward_demodulation, [], [f384, f1004])).
fof(f384, plain, ~ (op1(e13, e14) = op1(e14, e14)), inference(cnf_transformation, [], [f7])).
fof(f4252, plain, (~ spl68_56 | ~ spl68_66), inference(avatar_split_clause, [], [f4247, f1162, f1120])).
fof(f4247, plain, (~ (e10 = op1(e12, e13)) | ~ spl68_66), inference(backward_demodulation, [], [f410, f1164])).
fof(f410, plain, ~ (op1(e12, e11) = op1(e12, e13)), inference(cnf_transformation, [], [f7])).
fof(f4240, plain, (~ spl68_67 | ~ spl68_72), inference(avatar_split_clause, [], [f4234, f1187, f1166])).
fof(f1187, plain, (spl68_72 <=> (e11 = op1(e12, e10))), introduced(avatar_definition, [new_symbols(naming, [spl68_72])])).
fof(f4234, plain, (~ (e11 = op1(e12, e11)) | ~ spl68_72), inference(backward_demodulation, [], [f405, f1189])).
fof(f1189, plain, ((e11 = op1(e12, e10)) | ~ spl68_72), inference(avatar_component_clause, [], [f1187])).
fof(f4151, plain, (~ spl68_106 | ~ spl68_111), inference(avatar_split_clause, [], [f4145, f1351, f1330])).
fof(f4145, plain, (~ (e10 = op1(e10, e13)) | ~ spl68_111), inference(backward_demodulation, [], [f392, f1353])).
fof(f392, plain, ~ (op1(e10, e12) = op1(e10, e13)), inference(cnf_transformation, [], [f7])).
fof(f4148, plain, (~ spl68_61 | ~ spl68_111), inference(avatar_split_clause, [], [f4142, f1351, f1141])).
fof(f1141, plain, (spl68_61 <=> (e10 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl68_61])])).
fof(f4142, plain, (~ (e10 = op1(e12, e12)) | ~ spl68_111), inference(backward_demodulation, [], [f356, f1353])).
fof(f356, plain, ~ (op1(e10, e12) = op1(e12, e12)), inference(cnf_transformation, [], [f7])).
fof(f4087, plain, (~ spl68_131 | ~ spl68_132), inference(avatar_contradiction_clause, [], [f4086])).
fof(f4086, plain, ($false | (~ spl68_131 | ~ spl68_132)), inference(subsumption_resolution, [], [f4085, f545])).
fof(f4085, plain, ((e20 = e21) | (~ spl68_131 | ~ spl68_132)), inference(backward_demodulation, [], [f1491, f1487])).
fof(f1487, plain, ((e20 = op2(e24, e24)) | ~ spl68_131), inference(avatar_component_clause, [], [f1485])).
fof(f1485, plain, (spl68_131 <=> (e20 = op2(e24, e24))), introduced(avatar_definition, [new_symbols(naming, [spl68_131])])).
fof(f4084, plain, (~ spl68_132 | ~ spl68_134), inference(avatar_contradiction_clause, [], [f4083])).
fof(f4083, plain, ($false | (~ spl68_132 | ~ spl68_134)), inference(subsumption_resolution, [], [f4082, f550])).
fof(f4082, plain, ((e21 = e23) | (~ spl68_132 | ~ spl68_134)), inference(backward_demodulation, [], [f1499, f1491])).
fof(f1499, plain, ((e23 = op2(e24, e24)) | ~ spl68_134), inference(avatar_component_clause, [], [f1497])).
fof(f1497, plain, (spl68_134 <=> (e23 = op2(e24, e24))), introduced(avatar_definition, [new_symbols(naming, [spl68_134])])).
fof(f4058, plain, (~ spl68_140 | ~ spl68_145), inference(avatar_split_clause, [], [f4055, f1543, f1522])).
fof(f1522, plain, (spl68_140 <=> (e24 = op2(e24, e23))), introduced(avatar_definition, [new_symbols(naming, [spl68_140])])).
fof(f4055, plain, (~ (e24 = op2(e24, e23)) | ~ spl68_145), inference(backward_demodulation, [], [f532, f1545])).
fof(f532, plain, ~ (op2(e24, e22) = op2(e24, e23)), inference(cnf_transformation, [], [f8])).
fof(f4051, plain, (~ spl68_148 | ~ spl68_149), inference(avatar_contradiction_clause, [], [f4050])).
fof(f4050, plain, ($false | (~ spl68_148 | ~ spl68_149)), inference(subsumption_resolution, [], [f4049, f552])).
fof(f4049, plain, ((e22 = e23) | (~ spl68_148 | ~ spl68_149)), inference(backward_demodulation, [], [f1562, f1558])).
fof(f1558, plain, ((e22 = op2(e24, e21)) | ~ spl68_148), inference(avatar_component_clause, [], [f1556])).
fof(f1556, plain, (spl68_148 <=> (e22 = op2(e24, e21))), introduced(avatar_definition, [new_symbols(naming, [spl68_148])])).
fof(f4048, plain, (~ spl68_149 | ~ spl68_150), inference(avatar_contradiction_clause, [], [f4047])).
fof(f4047, plain, ($false | (~ spl68_149 | ~ spl68_150)), inference(subsumption_resolution, [], [f4046, f554])).
fof(f4046, plain, ((e23 = e24) | (~ spl68_149 | ~ spl68_150)), inference(backward_demodulation, [], [f1566, f1562])).
fof(f1566, plain, ((e24 = op2(e24, e21)) | ~ spl68_150), inference(avatar_component_clause, [], [f1564])).
fof(f1564, plain, (spl68_150 <=> (e24 = op2(e24, e21))), introduced(avatar_definition, [new_symbols(naming, [spl68_150])])).
fof(f4020, plain, (~ spl68_157 | ~ spl68_158), inference(avatar_contradiction_clause, [], [f4019])).
fof(f4019, plain, ($false | (~ spl68_157 | ~ spl68_158)), inference(subsumption_resolution, [], [f4018, f549])).
fof(f4018, plain, ((e21 = e22) | (~ spl68_157 | ~ spl68_158)), inference(backward_demodulation, [], [f1600, f1596])).
fof(f1596, plain, ((e21 = op2(e23, e24)) | ~ spl68_157), inference(avatar_component_clause, [], [f1594])).
fof(f1594, plain, (spl68_157 <=> (e21 = op2(e23, e24))), introduced(avatar_definition, [new_symbols(naming, [spl68_157])])).
fof(f4017, plain, (~ spl68_158 | ~ spl68_159), inference(avatar_contradiction_clause, [], [f4016])).
fof(f4016, plain, ($false | (~ spl68_158 | ~ spl68_159)), inference(subsumption_resolution, [], [f4015, f552])).
fof(f4015, plain, ((e22 = e23) | (~ spl68_158 | ~ spl68_159)), inference(backward_demodulation, [], [f1604, f1600])).
fof(f1604, plain, ((e23 = op2(e23, e24)) | ~ spl68_159), inference(avatar_component_clause, [], [f1602])).
fof(f1602, plain, (spl68_159 <=> (e23 = op2(e23, e24))), introduced(avatar_definition, [new_symbols(naming, [spl68_159])])).
fof(f4011, plain, (spl68_432 | ~ spl68_165), inference(avatar_split_clause, [], [f4010, f1627, f3224])).
fof(f4010, plain, ((e24 = h4(e11)) | ~ spl68_165), inference(backward_demodulation, [], [f755, f1629])).
fof(f4006, plain, (~ spl68_141 | ~ spl68_166), inference(avatar_split_clause, [], [f4002, f1632, f1527])).
fof(f1632, plain, (spl68_166 <=> (e20 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl68_166])])).
fof(f4002, plain, (~ (e20 = op2(e24, e22)) | ~ spl68_166), inference(backward_demodulation, [], [f464, f1634])).
fof(f1634, plain, ((e20 = op2(e23, e22)) | ~ spl68_166), inference(avatar_component_clause, [], [f1632])).
fof(f3958, plain, (spl68_346 | ~ spl68_191), inference(avatar_split_clause, [], [f3957, f1737, f2690])).
fof(f3957, plain, ((e20 = h3(e11)) | ~ spl68_191), inference(backward_demodulation, [], [f750, f1739])).
fof(f3951, plain, (~ spl68_171 | ~ spl68_196), inference(avatar_split_clause, [], [f3945, f1758, f1653])).
fof(f1758, plain, (spl68_196 <=> (e20 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl68_196])])).
fof(f3945, plain, (~ (e20 = op2(e23, e21)) | ~ spl68_196), inference(backward_demodulation, [], [f452, f1760])).
fof(f1760, plain, ((e20 = op2(e22, e21)) | ~ spl68_196), inference(avatar_component_clause, [], [f1758])).
fof(f3941, plain, (~ spl68_196 | ~ spl68_201), inference(avatar_split_clause, [], [f3935, f1779, f1758])).
fof(f3935, plain, (~ (e20 = op2(e22, e21)) | ~ spl68_201), inference(backward_demodulation, [], [f505, f1781])).
fof(f3940, plain, (~ spl68_151 | ~ spl68_201), inference(avatar_split_clause, [], [f3934, f1779, f1569])).
fof(f1569, plain, (spl68_151 <=> (e20 = op2(e24, e20))), introduced(avatar_definition, [new_symbols(naming, [spl68_151])])).
fof(f3934, plain, (~ (e20 = op2(e24, e20)) | ~ spl68_201), inference(backward_demodulation, [], [f443, f1781])).
fof(f443, plain, ~ (op2(e22, e20) = op2(e24, e20)), inference(cnf_transformation, [], [f8])).
fof(f3939, plain, (~ spl68_176 | ~ spl68_201), inference(avatar_split_clause, [], [f3933, f1779, f1674])).
fof(f3933, plain, (~ (e20 = op2(e23, e20)) | ~ spl68_201), inference(backward_demodulation, [], [f442, f1781])).
fof(f442, plain, ~ (op2(e22, e20) = op2(e23, e20)), inference(cnf_transformation, [], [f8])).
fof(f3913, plain, (~ spl68_206 | ~ spl68_211), inference(avatar_split_clause, [], [f3908, f1821, f1800])).
fof(f1821, plain, (spl68_211 <=> (e20 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl68_211])])).
fof(f3908, plain, (~ (e20 = op2(e21, e24)) | ~ spl68_211), inference(backward_demodulation, [], [f504, f1823])).
fof(f1823, plain, ((e20 = op2(e21, e23)) | ~ spl68_211), inference(avatar_component_clause, [], [f1821])).
fof(f504, plain, ~ (op2(e21, e23) = op2(e21, e24)), inference(cnf_transformation, [], [f8])).
fof(f3893, plain, (spl68_361 | ~ spl68_221), inference(avatar_split_clause, [], [f3892, f1863, f2766])).
fof(f1863, plain, (spl68_221 <=> (e20 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl68_221])])).
fof(f3892, plain, ((e20 = h2(e11)) | ~ spl68_221), inference(backward_demodulation, [], [f745, f1865])).
fof(f1865, plain, ((e20 = op2(e21, e21)) | ~ spl68_221), inference(avatar_component_clause, [], [f1863])).
fof(f3835, plain, (~ spl68_241 | ~ spl68_246), inference(avatar_split_clause, [], [f3828, f1968, f1947])).
fof(f1968, plain, (spl68_246 <=> (e20 = op2(e20, e21))), introduced(avatar_definition, [new_symbols(naming, [spl68_246])])).
fof(f3828, plain, (~ (e20 = op2(e20, e22)) | ~ spl68_246), inference(backward_demodulation, [], [f489, f1970])).
fof(f1970, plain, ((e20 = op2(e20, e21)) | ~ spl68_246), inference(avatar_component_clause, [], [f1968])).
fof(f3761, plain, (~ spl68_492 | ~ spl68_493 | ~ spl68_494 | ~ spl68_495 | ~ spl68_496 | ~ spl68_497 | ~ spl68_498 | ~ spl68_499 | ~ spl68_500 | ~ spl68_501 | ~ spl68_502 | ~ spl68_503 | ~ spl68_504 | ~ spl68_505 | ~ spl68_506 | ~ spl68_507 | spl68_373 | spl68_368 | spl68_363 | ~ spl68_519 | ~ spl68_509 | ~ spl68_510 | ~ spl68_511 | ~ spl68_512 | ~ spl68_513 | ~ spl68_514 | ~ spl68_515 | ~ spl68_516 | ~ spl68_517), inference(avatar_split_clause, [], [f3756, f3720, f3716, f3712, f3708, f3704, f3700, f3696, f3692, f3688, f3758, f2777, f2802, f2827, f3680, f3676, f3672, f3668, f3664, f3660, f3656, f3652, f3648, f3644, f3640, f3636, f3632, f3628, f3624, f3620])).
fof(f2827, plain, (spl68_373 <=> sP49), introduced(avatar_definition, [new_symbols(naming, [spl68_373])])).
fof(f2802, plain, (spl68_368 <=> sP50), introduced(avatar_definition, [new_symbols(naming, [spl68_368])])).
fof(f2777, plain, (spl68_363 <=> sP51), introduced(avatar_definition, [new_symbols(naming, [spl68_363])])).
fof(f3756, plain, (~ (h1(op1(e10, e14)) = op2(h1(e10), e20)) | ~ (h1(e10) = h1(op1(e11, e14))) | ~ (h1(op1(e12, e14)) = op2(h1(e12), e20)) | ~ (h1(e12) = h1(op1(e13, e14))) | ~ (h1(op1(e14, e10)) = op2(e20, h1(e10))) | ~ (h1(e13) = h1(op1(e14, e11))) | ~ (h1(op1(e14, e12)) = op2(e20, h1(e12))) | ~ (h1(op1(e14, e13)) = op2(e20, h1(e13))) | ~ (h1(e11) = h1(op1(e14, e14))) | ~ (e24 = h1(e11)) | sP51 | sP50 | sP49 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f3755, f738])).
fof(f3755, plain, (~ (h1(e10) = h1(op1(e11, e14))) | ~ (h1(op1(e12, e14)) = op2(h1(e12), e20)) | ~ (h1(e12) = h1(op1(e13, e14))) | ~ (h1(op1(e14, e10)) = op2(e20, h1(e10))) | ~ (h1(e13) = h1(op1(e14, e11))) | ~ (h1(op1(e14, e12)) = op2(e20, h1(e12))) | ~ (h1(op1(e14, e13)) = op2(e20, h1(e13))) | ~ (h1(e11) = h1(op1(e14, e14))) | ~ (e24 = h1(e11)) | sP51 | sP50 | sP49 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e14)) = op2(h1(e10), h1(e14))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f3754, f2438])).
fof(f3754, plain, (~ (h1(op1(e11, e14)) = op2(h1(e11), e20)) | ~ (h1(op1(e12, e14)) = op2(h1(e12), e20)) | ~ (h1(e12) = h1(op1(e13, e14))) | ~ (h1(op1(e14, e10)) = op2(e20, h1(e10))) | ~ (h1(e13) = h1(op1(e14, e11))) | ~ (h1(op1(e14, e12)) = op2(e20, h1(e12))) | ~ (h1(op1(e14, e13)) = op2(e20, h1(e13))) | ~ (h1(e11) = h1(op1(e14, e14))) | ~ (e24 = h1(e11)) | sP51 | sP50 | sP49 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e14)) = op2(h1(e10), h1(e14))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f3753, f738])).
fof(f3753, plain, (~ (h1(op1(e12, e14)) = op2(h1(e12), e20)) | ~ (h1(e12) = h1(op1(e13, e14))) | ~ (h1(op1(e14, e10)) = op2(e20, h1(e10))) | ~ (h1(e13) = h1(op1(e14, e11))) | ~ (h1(op1(e14, e12)) = op2(e20, h1(e12))) | ~ (h1(op1(e14, e13)) = op2(e20, h1(e13))) | ~ (h1(e11) = h1(op1(e14, e14))) | ~ (e24 = h1(e11)) | sP51 | sP50 | sP49 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e14)) = op2(h1(e11), h1(e14))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e14)) = op2(h1(e10), h1(e14))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f3752, f738])).
fof(f3752, plain, (~ (h1(e12) = h1(op1(e13, e14))) | ~ (h1(op1(e14, e10)) = op2(e20, h1(e10))) | ~ (h1(e13) = h1(op1(e14, e11))) | ~ (h1(op1(e14, e12)) = op2(e20, h1(e12))) | ~ (h1(op1(e14, e13)) = op2(e20, h1(e13))) | ~ (h1(e11) = h1(op1(e14, e14))) | ~ (e24 = h1(e11)) | sP51 | sP50 | sP49 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e14)) = op2(h1(e12), h1(e14))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e14)) = op2(h1(e11), h1(e14))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e14)) = op2(h1(e10), h1(e14))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f3751, f2428])).
fof(f3751, plain, (~ (h1(op1(e13, e14)) = op2(h1(e13), e20)) | ~ (h1(op1(e14, e10)) = op2(e20, h1(e10))) | ~ (h1(e13) = h1(op1(e14, e11))) | ~ (h1(op1(e14, e12)) = op2(e20, h1(e12))) | ~ (h1(op1(e14, e13)) = op2(e20, h1(e13))) | ~ (h1(e11) = h1(op1(e14, e14))) | ~ (e24 = h1(e11)) | sP51 | sP50 | sP49 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e14)) = op2(h1(e12), h1(e14))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e14)) = op2(h1(e11), h1(e14))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e14)) = op2(h1(e10), h1(e14))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f3750, f738])).
fof(f3750, plain, (~ (h1(op1(e14, e10)) = op2(e20, h1(e10))) | ~ (h1(e13) = h1(op1(e14, e11))) | ~ (h1(op1(e14, e12)) = op2(e20, h1(e12))) | ~ (h1(op1(e14, e13)) = op2(e20, h1(e13))) | ~ (h1(e11) = h1(op1(e14, e14))) | ~ (e24 = h1(e11)) | sP51 | sP50 | sP49 | ~ (h1(op1(e13, e14)) = op2(h1(e13), h1(e14))) | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e14)) = op2(h1(e12), h1(e14))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e14)) = op2(h1(e11), h1(e14))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e14)) = op2(h1(e10), h1(e14))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f3749, f738])).
fof(f3749, plain, (~ (h1(e13) = h1(op1(e14, e11))) | ~ (h1(op1(e14, e12)) = op2(e20, h1(e12))) | ~ (h1(op1(e14, e13)) = op2(e20, h1(e13))) | ~ (h1(e11) = h1(op1(e14, e14))) | ~ (e24 = h1(e11)) | sP51 | sP50 | sP49 | ~ (h1(op1(e14, e10)) = op2(h1(e14), h1(e10))) | ~ (h1(op1(e13, e14)) = op2(h1(e13), h1(e14))) | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e14)) = op2(h1(e12), h1(e14))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e14)) = op2(h1(e11), h1(e14))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e14)) = op2(h1(e10), h1(e14))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f3748, f2437])).
fof(f3748, plain, (~ (h1(op1(e14, e11)) = op2(e20, h1(e11))) | ~ (h1(op1(e14, e12)) = op2(e20, h1(e12))) | ~ (h1(op1(e14, e13)) = op2(e20, h1(e13))) | ~ (h1(e11) = h1(op1(e14, e14))) | ~ (e24 = h1(e11)) | sP51 | sP50 | sP49 | ~ (h1(op1(e14, e10)) = op2(h1(e14), h1(e10))) | ~ (h1(op1(e13, e14)) = op2(h1(e13), h1(e14))) | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e14)) = op2(h1(e12), h1(e14))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e14)) = op2(h1(e11), h1(e14))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e14)) = op2(h1(e10), h1(e14))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f3747, f738])).
fof(f3747, plain, (~ (h1(op1(e14, e12)) = op2(e20, h1(e12))) | ~ (h1(op1(e14, e13)) = op2(e20, h1(e13))) | ~ (h1(e11) = h1(op1(e14, e14))) | ~ (e24 = h1(e11)) | sP51 | sP50 | sP49 | ~ (h1(op1(e14, e11)) = op2(h1(e14), h1(e11))) | ~ (h1(op1(e14, e10)) = op2(h1(e14), h1(e10))) | ~ (h1(op1(e13, e14)) = op2(h1(e13), h1(e14))) | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e14)) = op2(h1(e12), h1(e14))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e14)) = op2(h1(e11), h1(e14))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e14)) = op2(h1(e10), h1(e14))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f3746, f738])).
fof(f3746, plain, (~ (h1(op1(e14, e13)) = op2(e20, h1(e13))) | ~ (h1(e11) = h1(op1(e14, e14))) | ~ (e24 = h1(e11)) | sP51 | sP50 | sP49 | ~ (h1(op1(e14, e12)) = op2(h1(e14), h1(e12))) | ~ (h1(op1(e14, e11)) = op2(h1(e14), h1(e11))) | ~ (h1(op1(e14, e10)) = op2(h1(e14), h1(e10))) | ~ (h1(op1(e13, e14)) = op2(h1(e13), h1(e14))) | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e14)) = op2(h1(e12), h1(e14))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e14)) = op2(h1(e11), h1(e14))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e14)) = op2(h1(e10), h1(e14))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f3745, f738])).
fof(f3745, plain, (~ (h1(e11) = h1(op1(e14, e14))) | ~ (e24 = h1(e11)) | sP51 | sP50 | sP49 | ~ (h1(op1(e14, e13)) = op2(h1(e14), h1(e13))) | ~ (h1(op1(e14, e12)) = op2(h1(e14), h1(e12))) | ~ (h1(op1(e14, e11)) = op2(h1(e14), h1(e11))) | ~ (h1(op1(e14, e10)) = op2(h1(e14), h1(e10))) | ~ (h1(op1(e13, e14)) = op2(h1(e13), h1(e14))) | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e14)) = op2(h1(e12), h1(e14))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e14)) = op2(h1(e11), h1(e14))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e14)) = op2(h1(e10), h1(e14))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f3744, f740])).
fof(f3744, plain, (~ (op2(e20, e20) = h1(op1(e14, e14))) | ~ (e24 = h1(e11)) | sP51 | sP50 | sP49 | ~ (h1(op1(e14, e13)) = op2(h1(e14), h1(e13))) | ~ (h1(op1(e14, e12)) = op2(h1(e14), h1(e12))) | ~ (h1(op1(e14, e11)) = op2(h1(e14), h1(e11))) | ~ (h1(op1(e14, e10)) = op2(h1(e14), h1(e10))) | ~ (h1(op1(e13, e14)) = op2(h1(e13), h1(e14))) | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e14)) = op2(h1(e12), h1(e14))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e14)) = op2(h1(e11), h1(e14))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e14)) = op2(h1(e10), h1(e14))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f3743, f738])).
fof(f3743, plain, (~ (e24 = h1(e11)) | sP51 | sP50 | sP49 | ~ (h1(op1(e14, e14)) = op2(h1(e14), h1(e14))) | ~ (h1(op1(e14, e13)) = op2(h1(e14), h1(e13))) | ~ (h1(op1(e14, e12)) = op2(h1(e14), h1(e12))) | ~ (h1(op1(e14, e11)) = op2(h1(e14), h1(e11))) | ~ (h1(op1(e14, e10)) = op2(h1(e14), h1(e10))) | ~ (h1(op1(e13, e14)) = op2(h1(e13), h1(e14))) | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e14)) = op2(h1(e12), h1(e14))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e14)) = op2(h1(e11), h1(e14))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e14)) = op2(h1(e10), h1(e14))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(subsumption_resolution, [], [f864, f2850])).
fof(f2850, plain, ~ sP48, inference(subsumption_resolution, [], [f862, f738])).
fof(f862, plain, (~ (e20 = h1(e14)) | ~ sP48), inference(cnf_transformation, [], [f162])).
fof(f162, plain, ((~ (e20 = h1(e14)) & ~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10))) | ~ sP48), inference(nnf_transformation, [], [f74])).
fof(f74, plain, ((~ (e20 = h1(e14)) & ~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10))) | ~ sP48), inference(usedef, [], [e74])).
fof(e74, plain, (sP48 <=> (~ (e20 = h1(e14)) & ~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP48])])).
fof(f864, plain, (~ (e24 = h1(e11)) | sP51 | sP50 | sP49 | sP48 | ~ (h1(op1(e14, e14)) = op2(h1(e14), h1(e14))) | ~ (h1(op1(e14, e13)) = op2(h1(e14), h1(e13))) | ~ (h1(op1(e14, e12)) = op2(h1(e14), h1(e12))) | ~ (h1(op1(e14, e11)) = op2(h1(e14), h1(e11))) | ~ (h1(op1(e14, e10)) = op2(h1(e14), h1(e10))) | ~ (h1(op1(e13, e14)) = op2(h1(e13), h1(e14))) | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e14)) = op2(h1(e12), h1(e14))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e14)) = op2(h1(e11), h1(e14))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e14)) = op2(h1(e10), h1(e14))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(cnf_transformation, [], [f94])).
fof(f94, plain, (((~ (e24 = h5(e14)) & ~ (e24 = h5(e13)) & ~ (e24 = h5(e12)) & ~ (e24 = h5(e11)) & ~ (e24 = h5(e10))) | sP67 | sP66 | sP65 | sP64 | ~ (h5(op1(e14, e14)) = op2(h5(e14), h5(e14))) | ~ (h5(op1(e14, e13)) = op2(h5(e14), h5(e13))) | ~ (h5(op1(e14, e12)) = op2(h5(e14), h5(e12))) | ~ (h5(op1(e14, e11)) = op2(h5(e14), h5(e11))) | ~ (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) | ~ (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))) & ((~ (e24 = h4(e14)) & ~ (e24 = h4(e13)) & ~ (e24 = h4(e12)) & ~ (e24 = h4(e11)) & ~ (e24 = h4(e10))) | sP63 | sP62 | sP61 | sP60 | ~ (h4(op1(e14, e14)) = op2(h4(e14), h4(e14))) | ~ (h4(op1(e14, e13)) = op2(h4(e14), h4(e13))) | ~ (h4(op1(e14, e12)) = op2(h4(e14), h4(e12))) | ~ (h4(op1(e14, e11)) = op2(h4(e14), h4(e11))) | ~ (h4(op1(e14, e10)) = op2(h4(e14), h4(e10))) | ~ (h4(op1(e13, e14)) = op2(h4(e13), h4(e14))) | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e14)) = op2(h4(e12), h4(e14))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e14)) = op2(h4(e11), h4(e14))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e14)) = op2(h4(e10), h4(e14))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) & ((~ (e24 = h3(e14)) & ~ (e24 = h3(e13)) & ~ (e24 = h3(e12)) & ~ (e24 = h3(e11)) & ~ (e24 = h3(e10))) | sP59 | sP58 | sP57 | sP56 | ~ (h3(op1(e14, e14)) = op2(h3(e14), h3(e14))) | ~ (h3(op1(e14, e13)) = op2(h3(e14), h3(e13))) | ~ (h3(op1(e14, e12)) = op2(h3(e14), h3(e12))) | ~ (h3(op1(e14, e11)) = op2(h3(e14), h3(e11))) | ~ (h3(op1(e14, e10)) = op2(h3(e14), h3(e10))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), h3(e14))) | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) & ((~ (e24 = h2(e14)) & ~ (e24 = h2(e13)) & ~ (e24 = h2(e12)) & ~ (e24 = h2(e11)) & ~ (e24 = h2(e10))) | sP55 | sP54 | sP53 | sP52 | ~ (h2(op1(e14, e14)) = op2(h2(e14), h2(e14))) | ~ (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) | ~ (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) | ~ (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) | ~ (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) | ~ (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) & ((~ (e24 = h1(e14)) & ~ (e24 = h1(e13)) & ~ (e24 = h1(e12)) & ~ (e24 = h1(e11)) & ~ (e24 = h1(e10))) | sP51 | sP50 | sP49 | sP48 | ~ (h1(op1(e14, e14)) = op2(h1(e14), h1(e14))) | ~ (h1(op1(e14, e13)) = op2(h1(e14), h1(e13))) | ~ (h1(op1(e14, e12)) = op2(h1(e14), h1(e12))) | ~ (h1(op1(e14, e11)) = op2(h1(e14), h1(e11))) | ~ (h1(op1(e14, e10)) = op2(h1(e14), h1(e10))) | ~ (h1(op1(e13, e14)) = op2(h1(e13), h1(e14))) | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e14)) = op2(h1(e12), h1(e14))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e14)) = op2(h1(e11), h1(e14))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e14)) = op2(h1(e10), h1(e14))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(definition_folding, [], [f23, e93, e92, e91, e90, e89, e88, e87, e86, e85, e84, e83, e82, e81, e80, e79, e78, e77, e76, e75, e74])).
fof(f75, plain, ((~ (e21 = h1(e14)) & ~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10))) | ~ sP49), inference(usedef, [], [e75])).
fof(e75, plain, (sP49 <=> (~ (e21 = h1(e14)) & ~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP49])])).
fof(f76, plain, ((~ (e22 = h1(e14)) & ~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10))) | ~ sP50), inference(usedef, [], [e76])).
fof(e76, plain, (sP50 <=> (~ (e22 = h1(e14)) & ~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP50])])).
fof(f77, plain, ((~ (e23 = h1(e14)) & ~ (e23 = h1(e13)) & ~ (e23 = h1(e12)) & ~ (e23 = h1(e11)) & ~ (e23 = h1(e10))) | ~ sP51), inference(usedef, [], [e77])).
fof(e77, plain, (sP51 <=> (~ (e23 = h1(e14)) & ~ (e23 = h1(e13)) & ~ (e23 = h1(e12)) & ~ (e23 = h1(e11)) & ~ (e23 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP51])])).
fof(f78, plain, ((~ (e20 = h2(e14)) & ~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10))) | ~ sP52), inference(usedef, [], [e78])).
fof(e78, plain, (sP52 <=> (~ (e20 = h2(e14)) & ~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP52])])).
fof(f79, plain, ((~ (e21 = h2(e14)) & ~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10))) | ~ sP53), inference(usedef, [], [e79])).
fof(e79, plain, (sP53 <=> (~ (e21 = h2(e14)) & ~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP53])])).
fof(f80, plain, ((~ (e22 = h2(e14)) & ~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10))) | ~ sP54), inference(usedef, [], [e80])).
fof(e80, plain, (sP54 <=> (~ (e22 = h2(e14)) & ~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP54])])).
fof(f81, plain, ((~ (e23 = h2(e14)) & ~ (e23 = h2(e13)) & ~ (e23 = h2(e12)) & ~ (e23 = h2(e11)) & ~ (e23 = h2(e10))) | ~ sP55), inference(usedef, [], [e81])).
fof(e81, plain, (sP55 <=> (~ (e23 = h2(e14)) & ~ (e23 = h2(e13)) & ~ (e23 = h2(e12)) & ~ (e23 = h2(e11)) & ~ (e23 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP55])])).
fof(f82, plain, ((~ (e20 = h3(e14)) & ~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10))) | ~ sP56), inference(usedef, [], [e82])).
fof(e82, plain, (sP56 <=> (~ (e20 = h3(e14)) & ~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP56])])).
fof(f83, plain, ((~ (e21 = h3(e14)) & ~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10))) | ~ sP57), inference(usedef, [], [e83])).
fof(e83, plain, (sP57 <=> (~ (e21 = h3(e14)) & ~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP57])])).
fof(f84, plain, ((~ (e22 = h3(e14)) & ~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10))) | ~ sP58), inference(usedef, [], [e84])).
fof(e84, plain, (sP58 <=> (~ (e22 = h3(e14)) & ~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP58])])).
fof(f85, plain, ((~ (e23 = h3(e14)) & ~ (e23 = h3(e13)) & ~ (e23 = h3(e12)) & ~ (e23 = h3(e11)) & ~ (e23 = h3(e10))) | ~ sP59), inference(usedef, [], [e85])).
fof(e85, plain, (sP59 <=> (~ (e23 = h3(e14)) & ~ (e23 = h3(e13)) & ~ (e23 = h3(e12)) & ~ (e23 = h3(e11)) & ~ (e23 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP59])])).
fof(f86, plain, ((~ (e20 = h4(e14)) & ~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10))) | ~ sP60), inference(usedef, [], [e86])).
fof(e86, plain, (sP60 <=> (~ (e20 = h4(e14)) & ~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP60])])).
fof(f87, plain, ((~ (e21 = h4(e14)) & ~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10))) | ~ sP61), inference(usedef, [], [e87])).
fof(e87, plain, (sP61 <=> (~ (e21 = h4(e14)) & ~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP61])])).
fof(f88, plain, ((~ (e22 = h4(e14)) & ~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10))) | ~ sP62), inference(usedef, [], [e88])).
fof(e88, plain, (sP62 <=> (~ (e22 = h4(e14)) & ~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP62])])).
fof(f89, plain, ((~ (e23 = h4(e14)) & ~ (e23 = h4(e13)) & ~ (e23 = h4(e12)) & ~ (e23 = h4(e11)) & ~ (e23 = h4(e10))) | ~ sP63), inference(usedef, [], [e89])).
fof(e89, plain, (sP63 <=> (~ (e23 = h4(e14)) & ~ (e23 = h4(e13)) & ~ (e23 = h4(e12)) & ~ (e23 = h4(e11)) & ~ (e23 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP63])])).
fof(f90, plain, ((~ (e20 = h5(e14)) & ~ (e20 = h5(e13)) & ~ (e20 = h5(e12)) & ~ (e20 = h5(e11)) & ~ (e20 = h5(e10))) | ~ sP64), inference(usedef, [], [e90])).
fof(e90, plain, (sP64 <=> (~ (e20 = h5(e14)) & ~ (e20 = h5(e13)) & ~ (e20 = h5(e12)) & ~ (e20 = h5(e11)) & ~ (e20 = h5(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP64])])).
fof(f91, plain, ((~ (e21 = h5(e14)) & ~ (e21 = h5(e13)) & ~ (e21 = h5(e12)) & ~ (e21 = h5(e11)) & ~ (e21 = h5(e10))) | ~ sP65), inference(usedef, [], [e91])).
fof(e91, plain, (sP65 <=> (~ (e21 = h5(e14)) & ~ (e21 = h5(e13)) & ~ (e21 = h5(e12)) & ~ (e21 = h5(e11)) & ~ (e21 = h5(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP65])])).
fof(f92, plain, ((~ (e22 = h5(e14)) & ~ (e22 = h5(e13)) & ~ (e22 = h5(e12)) & ~ (e22 = h5(e11)) & ~ (e22 = h5(e10))) | ~ sP66), inference(usedef, [], [e92])).
fof(e92, plain, (sP66 <=> (~ (e22 = h5(e14)) & ~ (e22 = h5(e13)) & ~ (e22 = h5(e12)) & ~ (e22 = h5(e11)) & ~ (e22 = h5(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP66])])).
fof(f93, plain, ((~ (e23 = h5(e14)) & ~ (e23 = h5(e13)) & ~ (e23 = h5(e12)) & ~ (e23 = h5(e11)) & ~ (e23 = h5(e10))) | ~ sP67), inference(usedef, [], [e93])).
fof(e93, plain, (sP67 <=> (~ (e23 = h5(e14)) & ~ (e23 = h5(e13)) & ~ (e23 = h5(e12)) & ~ (e23 = h5(e11)) & ~ (e23 = h5(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP67])])).
fof(f23, plain, (((~ (e24 = h5(e14)) & ~ (e24 = h5(e13)) & ~ (e24 = h5(e12)) & ~ (e24 = h5(e11)) & ~ (e24 = h5(e10))) | (~ (e23 = h5(e14)) & ~ (e23 = h5(e13)) & ~ (e23 = h5(e12)) & ~ (e23 = h5(e11)) & ~ (e23 = h5(e10))) | (~ (e22 = h5(e14)) & ~ (e22 = h5(e13)) & ~ (e22 = h5(e12)) & ~ (e22 = h5(e11)) & ~ (e22 = h5(e10))) | (~ (e21 = h5(e14)) & ~ (e21 = h5(e13)) & ~ (e21 = h5(e12)) & ~ (e21 = h5(e11)) & ~ (e21 = h5(e10))) | (~ (e20 = h5(e14)) & ~ (e20 = h5(e13)) & ~ (e20 = h5(e12)) & ~ (e20 = h5(e11)) & ~ (e20 = h5(e10))) | ~ (h5(op1(e14, e14)) = op2(h5(e14), h5(e14))) | ~ (h5(op1(e14, e13)) = op2(h5(e14), h5(e13))) | ~ (h5(op1(e14, e12)) = op2(h5(e14), h5(e12))) | ~ (h5(op1(e14, e11)) = op2(h5(e14), h5(e11))) | ~ (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) | ~ (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))) & ((~ (e24 = h4(e14)) & ~ (e24 = h4(e13)) & ~ (e24 = h4(e12)) & ~ (e24 = h4(e11)) & ~ (e24 = h4(e10))) | (~ (e23 = h4(e14)) & ~ (e23 = h4(e13)) & ~ (e23 = h4(e12)) & ~ (e23 = h4(e11)) & ~ (e23 = h4(e10))) | (~ (e22 = h4(e14)) & ~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10))) | (~ (e21 = h4(e14)) & ~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10))) | (~ (e20 = h4(e14)) & ~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10))) | ~ (h4(op1(e14, e14)) = op2(h4(e14), h4(e14))) | ~ (h4(op1(e14, e13)) = op2(h4(e14), h4(e13))) | ~ (h4(op1(e14, e12)) = op2(h4(e14), h4(e12))) | ~ (h4(op1(e14, e11)) = op2(h4(e14), h4(e11))) | ~ (h4(op1(e14, e10)) = op2(h4(e14), h4(e10))) | ~ (h4(op1(e13, e14)) = op2(h4(e13), h4(e14))) | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e14)) = op2(h4(e12), h4(e14))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e14)) = op2(h4(e11), h4(e14))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e14)) = op2(h4(e10), h4(e14))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) & ((~ (e24 = h3(e14)) & ~ (e24 = h3(e13)) & ~ (e24 = h3(e12)) & ~ (e24 = h3(e11)) & ~ (e24 = h3(e10))) | (~ (e23 = h3(e14)) & ~ (e23 = h3(e13)) & ~ (e23 = h3(e12)) & ~ (e23 = h3(e11)) & ~ (e23 = h3(e10))) | (~ (e22 = h3(e14)) & ~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10))) | (~ (e21 = h3(e14)) & ~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10))) | (~ (e20 = h3(e14)) & ~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10))) | ~ (h3(op1(e14, e14)) = op2(h3(e14), h3(e14))) | ~ (h3(op1(e14, e13)) = op2(h3(e14), h3(e13))) | ~ (h3(op1(e14, e12)) = op2(h3(e14), h3(e12))) | ~ (h3(op1(e14, e11)) = op2(h3(e14), h3(e11))) | ~ (h3(op1(e14, e10)) = op2(h3(e14), h3(e10))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), h3(e14))) | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) & ((~ (e24 = h2(e14)) & ~ (e24 = h2(e13)) & ~ (e24 = h2(e12)) & ~ (e24 = h2(e11)) & ~ (e24 = h2(e10))) | (~ (e23 = h2(e14)) & ~ (e23 = h2(e13)) & ~ (e23 = h2(e12)) & ~ (e23 = h2(e11)) & ~ (e23 = h2(e10))) | (~ (e22 = h2(e14)) & ~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10))) | (~ (e21 = h2(e14)) & ~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10))) | (~ (e20 = h2(e14)) & ~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10))) | ~ (h2(op1(e14, e14)) = op2(h2(e14), h2(e14))) | ~ (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) | ~ (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) | ~ (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) | ~ (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) | ~ (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) & ((~ (e24 = h1(e14)) & ~ (e24 = h1(e13)) & ~ (e24 = h1(e12)) & ~ (e24 = h1(e11)) & ~ (e24 = h1(e10))) | (~ (e23 = h1(e14)) & ~ (e23 = h1(e13)) & ~ (e23 = h1(e12)) & ~ (e23 = h1(e11)) & ~ (e23 = h1(e10))) | (~ (e22 = h1(e14)) & ~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10))) | (~ (e21 = h1(e14)) & ~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10))) | (~ (e20 = h1(e14)) & ~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10))) | ~ (h1(op1(e14, e14)) = op2(h1(e14), h1(e14))) | ~ (h1(op1(e14, e13)) = op2(h1(e14), h1(e13))) | ~ (h1(op1(e14, e12)) = op2(h1(e14), h1(e12))) | ~ (h1(op1(e14, e11)) = op2(h1(e14), h1(e11))) | ~ (h1(op1(e14, e10)) = op2(h1(e14), h1(e10))) | ~ (h1(op1(e13, e14)) = op2(h1(e13), h1(e14))) | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e14)) = op2(h1(e12), h1(e14))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e14)) = op2(h1(e11), h1(e14))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e14)) = op2(h1(e10), h1(e14))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(ennf_transformation, [], [f22])).
fof(f22, plain, ~ ((((e24 = h5(e14)) | (e24 = h5(e13)) | (e24 = h5(e12)) | (e24 = h5(e11)) | (e24 = h5(e10))) & ((e23 = h5(e14)) | (e23 = h5(e13)) | (e23 = h5(e12)) | (e23 = h5(e11)) | (e23 = h5(e10))) & ((e22 = h5(e14)) | (e22 = h5(e13)) | (e22 = h5(e12)) | (e22 = h5(e11)) | (e22 = h5(e10))) & ((e21 = h5(e14)) | (e21 = h5(e13)) | (e21 = h5(e12)) | (e21 = h5(e11)) | (e21 = h5(e10))) & ((e20 = h5(e14)) | (e20 = h5(e13)) | (e20 = h5(e12)) | (e20 = h5(e11)) | (e20 = h5(e10))) & (h5(op1(e14, e14)) = op2(h5(e14), h5(e14))) & (h5(op1(e14, e13)) = op2(h5(e14), h5(e13))) & (h5(op1(e14, e12)) = op2(h5(e14), h5(e12))) & (h5(op1(e14, e11)) = op2(h5(e14), h5(e11))) & (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) & (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) & (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) & (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) & (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) & (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) & (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) & (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) & (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) & (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) & (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) & (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) & (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) & (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) & (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) & (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) & (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) & (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) & (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) & (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) & (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))) | (((e24 = h4(e14)) | (e24 = h4(e13)) | (e24 = h4(e12)) | (e24 = h4(e11)) | (e24 = h4(e10))) & ((e23 = h4(e14)) | (e23 = h4(e13)) | (e23 = h4(e12)) | (e23 = h4(e11)) | (e23 = h4(e10))) & ((e22 = h4(e14)) | (e22 = h4(e13)) | (e22 = h4(e12)) | (e22 = h4(e11)) | (e22 = h4(e10))) & ((e21 = h4(e14)) | (e21 = h4(e13)) | (e21 = h4(e12)) | (e21 = h4(e11)) | (e21 = h4(e10))) & ((e20 = h4(e14)) | (e20 = h4(e13)) | (e20 = h4(e12)) | (e20 = h4(e11)) | (e20 = h4(e10))) & (h4(op1(e14, e14)) = op2(h4(e14), h4(e14))) & (h4(op1(e14, e13)) = op2(h4(e14), h4(e13))) & (h4(op1(e14, e12)) = op2(h4(e14), h4(e12))) & (h4(op1(e14, e11)) = op2(h4(e14), h4(e11))) & (h4(op1(e14, e10)) = op2(h4(e14), h4(e10))) & (h4(op1(e13, e14)) = op2(h4(e13), h4(e14))) & (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) & (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) & (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) & (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) & (h4(op1(e12, e14)) = op2(h4(e12), h4(e14))) & (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) & (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) & (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) & (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) & (h4(op1(e11, e14)) = op2(h4(e11), h4(e14))) & (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) & (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) & (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) & (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) & (h4(op1(e10, e14)) = op2(h4(e10), h4(e14))) & (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) & (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) & (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) & (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) | (((e24 = h3(e14)) | (e24 = h3(e13)) | (e24 = h3(e12)) | (e24 = h3(e11)) | (e24 = h3(e10))) & ((e23 = h3(e14)) | (e23 = h3(e13)) | (e23 = h3(e12)) | (e23 = h3(e11)) | (e23 = h3(e10))) & ((e22 = h3(e14)) | (e22 = h3(e13)) | (e22 = h3(e12)) | (e22 = h3(e11)) | (e22 = h3(e10))) & ((e21 = h3(e14)) | (e21 = h3(e13)) | (e21 = h3(e12)) | (e21 = h3(e11)) | (e21 = h3(e10))) & ((e20 = h3(e14)) | (e20 = h3(e13)) | (e20 = h3(e12)) | (e20 = h3(e11)) | (e20 = h3(e10))) & (h3(op1(e14, e14)) = op2(h3(e14), h3(e14))) & (h3(op1(e14, e13)) = op2(h3(e14), h3(e13))) & (h3(op1(e14, e12)) = op2(h3(e14), h3(e12))) & (h3(op1(e14, e11)) = op2(h3(e14), h3(e11))) & (h3(op1(e14, e10)) = op2(h3(e14), h3(e10))) & (h3(op1(e13, e14)) = op2(h3(e13), h3(e14))) & (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) & (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) & (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) & (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) & (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) & (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) & (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) & (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) & (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) & (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) & (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) & (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) & (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) & (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) & (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) & (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) & (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) & (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) & (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) | (((e24 = h2(e14)) | (e24 = h2(e13)) | (e24 = h2(e12)) | (e24 = h2(e11)) | (e24 = h2(e10))) & ((e23 = h2(e14)) | (e23 = h2(e13)) | (e23 = h2(e12)) | (e23 = h2(e11)) | (e23 = h2(e10))) & ((e22 = h2(e14)) | (e22 = h2(e13)) | (e22 = h2(e12)) | (e22 = h2(e11)) | (e22 = h2(e10))) & ((e21 = h2(e14)) | (e21 = h2(e13)) | (e21 = h2(e12)) | (e21 = h2(e11)) | (e21 = h2(e10))) & ((e20 = h2(e14)) | (e20 = h2(e13)) | (e20 = h2(e12)) | (e20 = h2(e11)) | (e20 = h2(e10))) & (h2(op1(e14, e14)) = op2(h2(e14), h2(e14))) & (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) & (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) & (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) & (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) & (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) & (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) & (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) & (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) & (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) & (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) & (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) & (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) & (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) & (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) & (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) & (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) & (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) & (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) & (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) & (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) & (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) & (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) & (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) & (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) | (((e24 = h1(e14)) | (e24 = h1(e13)) | (e24 = h1(e12)) | (e24 = h1(e11)) | (e24 = h1(e10))) & ((e23 = h1(e14)) | (e23 = h1(e13)) | (e23 = h1(e12)) | (e23 = h1(e11)) | (e23 = h1(e10))) & ((e22 = h1(e14)) | (e22 = h1(e13)) | (e22 = h1(e12)) | (e22 = h1(e11)) | (e22 = h1(e10))) & ((e21 = h1(e14)) | (e21 = h1(e13)) | (e21 = h1(e12)) | (e21 = h1(e11)) | (e21 = h1(e10))) & ((e20 = h1(e14)) | (e20 = h1(e13)) | (e20 = h1(e12)) | (e20 = h1(e11)) | (e20 = h1(e10))) & (h1(op1(e14, e14)) = op2(h1(e14), h1(e14))) & (h1(op1(e14, e13)) = op2(h1(e14), h1(e13))) & (h1(op1(e14, e12)) = op2(h1(e14), h1(e12))) & (h1(op1(e14, e11)) = op2(h1(e14), h1(e11))) & (h1(op1(e14, e10)) = op2(h1(e14), h1(e10))) & (h1(op1(e13, e14)) = op2(h1(e13), h1(e14))) & (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) & (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) & (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) & (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) & (h1(op1(e12, e14)) = op2(h1(e12), h1(e14))) & (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) & (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) & (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) & (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) & (h1(op1(e11, e14)) = op2(h1(e11), h1(e14))) & (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) & (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) & (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) & (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) & (h1(op1(e10, e14)) = op2(h1(e10), h1(e14))) & (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) & (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) & (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) & (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(negated_conjecture, [], [f21])).
fof(f21, plain, ~ ((((e24 = h5(e14)) | (e24 = h5(e13)) | (e24 = h5(e12)) | (e24 = h5(e11)) | (e24 = h5(e10))) & ((e23 = h5(e14)) | (e23 = h5(e13)) | (e23 = h5(e12)) | (e23 = h5(e11)) | (e23 = h5(e10))) & ((e22 = h5(e14)) | (e22 = h5(e13)) | (e22 = h5(e12)) | (e22 = h5(e11)) | (e22 = h5(e10))) & ((e21 = h5(e14)) | (e21 = h5(e13)) | (e21 = h5(e12)) | (e21 = h5(e11)) | (e21 = h5(e10))) & ((e20 = h5(e14)) | (e20 = h5(e13)) | (e20 = h5(e12)) | (e20 = h5(e11)) | (e20 = h5(e10))) & (h5(op1(e14, e14)) = op2(h5(e14), h5(e14))) & (h5(op1(e14, e13)) = op2(h5(e14), h5(e13))) & (h5(op1(e14, e12)) = op2(h5(e14), h5(e12))) & (h5(op1(e14, e11)) = op2(h5(e14), h5(e11))) & (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) & (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) & (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) & (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) & (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) & (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) & (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) & (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) & (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) & (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) & (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) & (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) & (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) & (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) & (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) & (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) & (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) & (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) & (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) & (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) & (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))) | (((e24 = h4(e14)) | (e24 = h4(e13)) | (e24 = h4(e12)) | (e24 = h4(e11)) | (e24 = h4(e10))) & ((e23 = h4(e14)) | (e23 = h4(e13)) | (e23 = h4(e12)) | (e23 = h4(e11)) | (e23 = h4(e10))) & ((e22 = h4(e14)) | (e22 = h4(e13)) | (e22 = h4(e12)) | (e22 = h4(e11)) | (e22 = h4(e10))) & ((e21 = h4(e14)) | (e21 = h4(e13)) | (e21 = h4(e12)) | (e21 = h4(e11)) | (e21 = h4(e10))) & ((e20 = h4(e14)) | (e20 = h4(e13)) | (e20 = h4(e12)) | (e20 = h4(e11)) | (e20 = h4(e10))) & (h4(op1(e14, e14)) = op2(h4(e14), h4(e14))) & (h4(op1(e14, e13)) = op2(h4(e14), h4(e13))) & (h4(op1(e14, e12)) = op2(h4(e14), h4(e12))) & (h4(op1(e14, e11)) = op2(h4(e14), h4(e11))) & (h4(op1(e14, e10)) = op2(h4(e14), h4(e10))) & (h4(op1(e13, e14)) = op2(h4(e13), h4(e14))) & (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) & (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) & (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) & (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) & (h4(op1(e12, e14)) = op2(h4(e12), h4(e14))) & (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) & (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) & (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) & (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) & (h4(op1(e11, e14)) = op2(h4(e11), h4(e14))) & (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) & (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) & (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) & (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) & (h4(op1(e10, e14)) = op2(h4(e10), h4(e14))) & (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) & (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) & (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) & (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) | (((e24 = h3(e14)) | (e24 = h3(e13)) | (e24 = h3(e12)) | (e24 = h3(e11)) | (e24 = h3(e10))) & ((e23 = h3(e14)) | (e23 = h3(e13)) | (e23 = h3(e12)) | (e23 = h3(e11)) | (e23 = h3(e10))) & ((e22 = h3(e14)) | (e22 = h3(e13)) | (e22 = h3(e12)) | (e22 = h3(e11)) | (e22 = h3(e10))) & ((e21 = h3(e14)) | (e21 = h3(e13)) | (e21 = h3(e12)) | (e21 = h3(e11)) | (e21 = h3(e10))) & ((e20 = h3(e14)) | (e20 = h3(e13)) | (e20 = h3(e12)) | (e20 = h3(e11)) | (e20 = h3(e10))) & (h3(op1(e14, e14)) = op2(h3(e14), h3(e14))) & (h3(op1(e14, e13)) = op2(h3(e14), h3(e13))) & (h3(op1(e14, e12)) = op2(h3(e14), h3(e12))) & (h3(op1(e14, e11)) = op2(h3(e14), h3(e11))) & (h3(op1(e14, e10)) = op2(h3(e14), h3(e10))) & (h3(op1(e13, e14)) = op2(h3(e13), h3(e14))) & (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) & (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) & (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) & (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) & (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) & (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) & (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) & (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) & (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) & (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) & (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) & (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) & (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) & (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) & (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) & (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) & (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) & (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) & (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) | (((e24 = h2(e14)) | (e24 = h2(e13)) | (e24 = h2(e12)) | (e24 = h2(e11)) | (e24 = h2(e10))) & ((e23 = h2(e14)) | (e23 = h2(e13)) | (e23 = h2(e12)) | (e23 = h2(e11)) | (e23 = h2(e10))) & ((e22 = h2(e14)) | (e22 = h2(e13)) | (e22 = h2(e12)) | (e22 = h2(e11)) | (e22 = h2(e10))) & ((e21 = h2(e14)) | (e21 = h2(e13)) | (e21 = h2(e12)) | (e21 = h2(e11)) | (e21 = h2(e10))) & ((e20 = h2(e14)) | (e20 = h2(e13)) | (e20 = h2(e12)) | (e20 = h2(e11)) | (e20 = h2(e10))) & (h2(op1(e14, e14)) = op2(h2(e14), h2(e14))) & (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) & (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) & (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) & (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) & (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) & (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) & (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) & (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) & (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) & (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) & (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) & (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) & (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) & (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) & (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) & (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) & (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) & (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) & (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) & (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) & (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) & (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) & (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) & (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) | (((e24 = h1(e14)) | (e24 = h1(e13)) | (e24 = h1(e12)) | (e24 = h1(e11)) | (e24 = h1(e10))) & ((e23 = h1(e14)) | (e23 = h1(e13)) | (e23 = h1(e12)) | (e23 = h1(e11)) | (e23 = h1(e10))) & ((e22 = h1(e14)) | (e22 = h1(e13)) | (e22 = h1(e12)) | (e22 = h1(e11)) | (e22 = h1(e10))) & ((e21 = h1(e14)) | (e21 = h1(e13)) | (e21 = h1(e12)) | (e21 = h1(e11)) | (e21 = h1(e10))) & ((e20 = h1(e14)) | (e20 = h1(e13)) | (e20 = h1(e12)) | (e20 = h1(e11)) | (e20 = h1(e10))) & (h1(op1(e14, e14)) = op2(h1(e14), h1(e14))) & (h1(op1(e14, e13)) = op2(h1(e14), h1(e13))) & (h1(op1(e14, e12)) = op2(h1(e14), h1(e12))) & (h1(op1(e14, e11)) = op2(h1(e14), h1(e11))) & (h1(op1(e14, e10)) = op2(h1(e14), h1(e10))) & (h1(op1(e13, e14)) = op2(h1(e13), h1(e14))) & (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) & (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) & (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) & (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) & (h1(op1(e12, e14)) = op2(h1(e12), h1(e14))) & (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) & (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) & (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) & (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) & (h1(op1(e11, e14)) = op2(h1(e11), h1(e14))) & (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) & (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) & (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) & (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) & (h1(op1(e10, e14)) = op2(h1(e10), h1(e14))) & (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) & (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) & (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) & (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG052+1.p', co1)).
fof(f3602, plain, (~ spl68_463 | ~ spl68_464 | ~ spl68_465 | ~ spl68_466 | ~ spl68_467 | ~ spl68_468 | ~ spl68_469 | ~ spl68_470 | ~ spl68_471 | ~ spl68_472 | ~ spl68_473 | ~ spl68_474 | ~ spl68_475 | ~ spl68_476 | ~ spl68_477 | ~ spl68_478 | spl68_358 | spl68_353 | spl68_348 | ~ spl68_491 | ~ spl68_480 | ~ spl68_481 | ~ spl68_482 | ~ spl68_483 | ~ spl68_484 | ~ spl68_485 | ~ spl68_486 | ~ spl68_487 | ~ spl68_488), inference(avatar_split_clause, [], [f3597, f3542, f3538, f3534, f3530, f3526, f3522, f3518, f3514, f3510, f3599, f2701, f2726, f2752, f3502, f3498, f3494, f3490, f3486, f3482, f3478, f3474, f3470, f3466, f3462, f3458, f3454, f3450, f3446, f3442])).
fof(f2752, plain, (spl68_358 <=> sP52), introduced(avatar_definition, [new_symbols(naming, [spl68_358])])).
fof(f2726, plain, (spl68_353 <=> sP54), introduced(avatar_definition, [new_symbols(naming, [spl68_353])])).
fof(f2701, plain, (spl68_348 <=> sP55), introduced(avatar_definition, [new_symbols(naming, [spl68_348])])).
fof(f3597, plain, (~ (h2(op1(e10, e14)) = op2(h2(e10), e21)) | ~ (h2(e10) = h2(op1(e11, e14))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), e21)) | ~ (h2(e12) = h2(op1(e13, e14))) | ~ (h2(op1(e14, e10)) = op2(e21, h2(e10))) | ~ (h2(e13) = h2(op1(e14, e11))) | ~ (h2(op1(e14, e12)) = op2(e21, h2(e12))) | ~ (h2(op1(e14, e13)) = op2(e21, h2(e13))) | ~ (h2(e11) = h2(op1(e14, e14))) | ~ (e24 = h2(e10)) | sP55 | sP54 | sP52 | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f3596, f743])).
fof(f3596, plain, (~ (h2(e10) = h2(op1(e11, e14))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), e21)) | ~ (h2(e12) = h2(op1(e13, e14))) | ~ (h2(op1(e14, e10)) = op2(e21, h2(e10))) | ~ (h2(e13) = h2(op1(e14, e11))) | ~ (h2(op1(e14, e12)) = op2(e21, h2(e12))) | ~ (h2(op1(e14, e13)) = op2(e21, h2(e13))) | ~ (h2(e11) = h2(op1(e14, e14))) | ~ (e24 = h2(e10)) | sP55 | sP54 | sP52 | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f3595, f2449])).
fof(f3595, plain, (~ (h2(op1(e11, e14)) = op2(h2(e11), e21)) | ~ (h2(op1(e12, e14)) = op2(h2(e12), e21)) | ~ (h2(e12) = h2(op1(e13, e14))) | ~ (h2(op1(e14, e10)) = op2(e21, h2(e10))) | ~ (h2(e13) = h2(op1(e14, e11))) | ~ (h2(op1(e14, e12)) = op2(e21, h2(e12))) | ~ (h2(op1(e14, e13)) = op2(e21, h2(e13))) | ~ (h2(e11) = h2(op1(e14, e14))) | ~ (e24 = h2(e10)) | sP55 | sP54 | sP52 | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f3594, f743])).
fof(f3594, plain, (~ (h2(op1(e12, e14)) = op2(h2(e12), e21)) | ~ (h2(e12) = h2(op1(e13, e14))) | ~ (h2(op1(e14, e10)) = op2(e21, h2(e10))) | ~ (h2(e13) = h2(op1(e14, e11))) | ~ (h2(op1(e14, e12)) = op2(e21, h2(e12))) | ~ (h2(op1(e14, e13)) = op2(e21, h2(e13))) | ~ (h2(e11) = h2(op1(e14, e14))) | ~ (e24 = h2(e10)) | sP55 | sP54 | sP52 | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f3593, f743])).
fof(f3593, plain, (~ (h2(e12) = h2(op1(e13, e14))) | ~ (h2(op1(e14, e10)) = op2(e21, h2(e10))) | ~ (h2(e13) = h2(op1(e14, e11))) | ~ (h2(op1(e14, e12)) = op2(e21, h2(e12))) | ~ (h2(op1(e14, e13)) = op2(e21, h2(e13))) | ~ (h2(e11) = h2(op1(e14, e14))) | ~ (e24 = h2(e10)) | sP55 | sP54 | sP52 | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f3592, f2439])).
fof(f3592, plain, (~ (h2(op1(e13, e14)) = op2(h2(e13), e21)) | ~ (h2(op1(e14, e10)) = op2(e21, h2(e10))) | ~ (h2(e13) = h2(op1(e14, e11))) | ~ (h2(op1(e14, e12)) = op2(e21, h2(e12))) | ~ (h2(op1(e14, e13)) = op2(e21, h2(e13))) | ~ (h2(e11) = h2(op1(e14, e14))) | ~ (e24 = h2(e10)) | sP55 | sP54 | sP52 | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f3591, f743])).
fof(f3591, plain, (~ (h2(op1(e14, e10)) = op2(e21, h2(e10))) | ~ (h2(e13) = h2(op1(e14, e11))) | ~ (h2(op1(e14, e12)) = op2(e21, h2(e12))) | ~ (h2(op1(e14, e13)) = op2(e21, h2(e13))) | ~ (h2(e11) = h2(op1(e14, e14))) | ~ (e24 = h2(e10)) | sP55 | sP54 | sP52 | ~ (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f3590, f743])).
fof(f3590, plain, (~ (h2(e13) = h2(op1(e14, e11))) | ~ (h2(op1(e14, e12)) = op2(e21, h2(e12))) | ~ (h2(op1(e14, e13)) = op2(e21, h2(e13))) | ~ (h2(e11) = h2(op1(e14, e14))) | ~ (e24 = h2(e10)) | sP55 | sP54 | sP52 | ~ (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) | ~ (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f3589, f2448])).
fof(f3589, plain, (~ (h2(op1(e14, e11)) = op2(e21, h2(e11))) | ~ (h2(op1(e14, e12)) = op2(e21, h2(e12))) | ~ (h2(op1(e14, e13)) = op2(e21, h2(e13))) | ~ (h2(e11) = h2(op1(e14, e14))) | ~ (e24 = h2(e10)) | sP55 | sP54 | sP52 | ~ (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) | ~ (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f3588, f743])).
fof(f3588, plain, (~ (h2(op1(e14, e12)) = op2(e21, h2(e12))) | ~ (h2(op1(e14, e13)) = op2(e21, h2(e13))) | ~ (h2(e11) = h2(op1(e14, e14))) | ~ (e24 = h2(e10)) | sP55 | sP54 | sP52 | ~ (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) | ~ (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) | ~ (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f3587, f743])).
fof(f3587, plain, (~ (h2(op1(e14, e13)) = op2(e21, h2(e13))) | ~ (h2(e11) = h2(op1(e14, e14))) | ~ (e24 = h2(e10)) | sP55 | sP54 | sP52 | ~ (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) | ~ (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) | ~ (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) | ~ (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f3586, f743])).
fof(f3586, plain, (~ (h2(e11) = h2(op1(e14, e14))) | ~ (e24 = h2(e10)) | sP55 | sP54 | sP52 | ~ (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) | ~ (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) | ~ (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) | ~ (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) | ~ (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f3585, f745])).
fof(f3585, plain, (~ (op2(e21, e21) = h2(op1(e14, e14))) | ~ (e24 = h2(e10)) | sP55 | sP54 | sP52 | ~ (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) | ~ (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) | ~ (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) | ~ (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) | ~ (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f3584, f743])).
fof(f3584, plain, (~ (e24 = h2(e10)) | sP55 | sP54 | sP52 | ~ (h2(op1(e14, e14)) = op2(h2(e14), h2(e14))) | ~ (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) | ~ (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) | ~ (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) | ~ (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) | ~ (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(subsumption_resolution, [], [f868, f2749])).
fof(f2749, plain, ~ sP53, inference(subsumption_resolution, [], [f837, f743])).
fof(f837, plain, (~ (e21 = h2(e14)) | ~ sP53), inference(cnf_transformation, [], [f157])).
fof(f157, plain, ((~ (e21 = h2(e14)) & ~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10))) | ~ sP53), inference(nnf_transformation, [], [f79])).
fof(f868, plain, (~ (e24 = h2(e10)) | sP55 | sP54 | sP53 | sP52 | ~ (h2(op1(e14, e14)) = op2(h2(e14), h2(e14))) | ~ (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) | ~ (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) | ~ (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) | ~ (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) | ~ (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(cnf_transformation, [], [f94])).
fof(f3583, plain, (~ spl68_463 | ~ spl68_464 | ~ spl68_465 | ~ spl68_466 | ~ spl68_467 | ~ spl68_468 | ~ spl68_469 | ~ spl68_470 | ~ spl68_471 | ~ spl68_472 | ~ spl68_473 | ~ spl68_474 | ~ spl68_475 | ~ spl68_476 | ~ spl68_477 | ~ spl68_478 | spl68_358 | spl68_353 | spl68_348 | ~ spl68_490 | ~ spl68_480 | ~ spl68_481 | ~ spl68_482 | ~ spl68_483 | ~ spl68_484 | ~ spl68_485 | ~ spl68_486 | ~ spl68_487 | ~ spl68_488), inference(avatar_split_clause, [], [f3578, f3542, f3538, f3534, f3530, f3526, f3522, f3518, f3514, f3510, f3580, f2701, f2726, f2752, f3502, f3498, f3494, f3490, f3486, f3482, f3478, f3474, f3470, f3466, f3462, f3458, f3454, f3450, f3446, f3442])).
fof(f3578, plain, (~ (h2(op1(e10, e14)) = op2(h2(e10), e21)) | ~ (h2(e10) = h2(op1(e11, e14))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), e21)) | ~ (h2(e12) = h2(op1(e13, e14))) | ~ (h2(op1(e14, e10)) = op2(e21, h2(e10))) | ~ (h2(e13) = h2(op1(e14, e11))) | ~ (h2(op1(e14, e12)) = op2(e21, h2(e12))) | ~ (h2(op1(e14, e13)) = op2(e21, h2(e13))) | ~ (h2(e11) = h2(op1(e14, e14))) | ~ (e24 = h2(e11)) | sP55 | sP54 | sP52 | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f3577, f743])).
fof(f3577, plain, (~ (h2(e10) = h2(op1(e11, e14))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), e21)) | ~ (h2(e12) = h2(op1(e13, e14))) | ~ (h2(op1(e14, e10)) = op2(e21, h2(e10))) | ~ (h2(e13) = h2(op1(e14, e11))) | ~ (h2(op1(e14, e12)) = op2(e21, h2(e12))) | ~ (h2(op1(e14, e13)) = op2(e21, h2(e13))) | ~ (h2(e11) = h2(op1(e14, e14))) | ~ (e24 = h2(e11)) | sP55 | sP54 | sP52 | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f3576, f2449])).
fof(f3576, plain, (~ (h2(op1(e11, e14)) = op2(h2(e11), e21)) | ~ (h2(op1(e12, e14)) = op2(h2(e12), e21)) | ~ (h2(e12) = h2(op1(e13, e14))) | ~ (h2(op1(e14, e10)) = op2(e21, h2(e10))) | ~ (h2(e13) = h2(op1(e14, e11))) | ~ (h2(op1(e14, e12)) = op2(e21, h2(e12))) | ~ (h2(op1(e14, e13)) = op2(e21, h2(e13))) | ~ (h2(e11) = h2(op1(e14, e14))) | ~ (e24 = h2(e11)) | sP55 | sP54 | sP52 | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f3575, f743])).
fof(f3575, plain, (~ (h2(op1(e12, e14)) = op2(h2(e12), e21)) | ~ (h2(e12) = h2(op1(e13, e14))) | ~ (h2(op1(e14, e10)) = op2(e21, h2(e10))) | ~ (h2(e13) = h2(op1(e14, e11))) | ~ (h2(op1(e14, e12)) = op2(e21, h2(e12))) | ~ (h2(op1(e14, e13)) = op2(e21, h2(e13))) | ~ (h2(e11) = h2(op1(e14, e14))) | ~ (e24 = h2(e11)) | sP55 | sP54 | sP52 | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f3574, f743])).
fof(f3574, plain, (~ (h2(e12) = h2(op1(e13, e14))) | ~ (h2(op1(e14, e10)) = op2(e21, h2(e10))) | ~ (h2(e13) = h2(op1(e14, e11))) | ~ (h2(op1(e14, e12)) = op2(e21, h2(e12))) | ~ (h2(op1(e14, e13)) = op2(e21, h2(e13))) | ~ (h2(e11) = h2(op1(e14, e14))) | ~ (e24 = h2(e11)) | sP55 | sP54 | sP52 | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f3573, f2439])).
fof(f3573, plain, (~ (h2(op1(e13, e14)) = op2(h2(e13), e21)) | ~ (h2(op1(e14, e10)) = op2(e21, h2(e10))) | ~ (h2(e13) = h2(op1(e14, e11))) | ~ (h2(op1(e14, e12)) = op2(e21, h2(e12))) | ~ (h2(op1(e14, e13)) = op2(e21, h2(e13))) | ~ (h2(e11) = h2(op1(e14, e14))) | ~ (e24 = h2(e11)) | sP55 | sP54 | sP52 | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f3572, f743])).
fof(f3572, plain, (~ (h2(op1(e14, e10)) = op2(e21, h2(e10))) | ~ (h2(e13) = h2(op1(e14, e11))) | ~ (h2(op1(e14, e12)) = op2(e21, h2(e12))) | ~ (h2(op1(e14, e13)) = op2(e21, h2(e13))) | ~ (h2(e11) = h2(op1(e14, e14))) | ~ (e24 = h2(e11)) | sP55 | sP54 | sP52 | ~ (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f3571, f743])).
fof(f3571, plain, (~ (h2(e13) = h2(op1(e14, e11))) | ~ (h2(op1(e14, e12)) = op2(e21, h2(e12))) | ~ (h2(op1(e14, e13)) = op2(e21, h2(e13))) | ~ (h2(e11) = h2(op1(e14, e14))) | ~ (e24 = h2(e11)) | sP55 | sP54 | sP52 | ~ (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) | ~ (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f3570, f2448])).
fof(f3570, plain, (~ (h2(op1(e14, e11)) = op2(e21, h2(e11))) | ~ (h2(op1(e14, e12)) = op2(e21, h2(e12))) | ~ (h2(op1(e14, e13)) = op2(e21, h2(e13))) | ~ (h2(e11) = h2(op1(e14, e14))) | ~ (e24 = h2(e11)) | sP55 | sP54 | sP52 | ~ (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) | ~ (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f3569, f743])).
fof(f3569, plain, (~ (h2(op1(e14, e12)) = op2(e21, h2(e12))) | ~ (h2(op1(e14, e13)) = op2(e21, h2(e13))) | ~ (h2(e11) = h2(op1(e14, e14))) | ~ (e24 = h2(e11)) | sP55 | sP54 | sP52 | ~ (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) | ~ (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) | ~ (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f3568, f743])).
fof(f3568, plain, (~ (h2(op1(e14, e13)) = op2(e21, h2(e13))) | ~ (h2(e11) = h2(op1(e14, e14))) | ~ (e24 = h2(e11)) | sP55 | sP54 | sP52 | ~ (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) | ~ (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) | ~ (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) | ~ (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f3567, f743])).
fof(f3567, plain, (~ (h2(e11) = h2(op1(e14, e14))) | ~ (e24 = h2(e11)) | sP55 | sP54 | sP52 | ~ (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) | ~ (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) | ~ (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) | ~ (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) | ~ (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f3566, f745])).
fof(f3566, plain, (~ (op2(e21, e21) = h2(op1(e14, e14))) | ~ (e24 = h2(e11)) | sP55 | sP54 | sP52 | ~ (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) | ~ (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) | ~ (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) | ~ (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) | ~ (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f3565, f743])).
fof(f3565, plain, (~ (e24 = h2(e11)) | sP55 | sP54 | sP52 | ~ (h2(op1(e14, e14)) = op2(h2(e14), h2(e14))) | ~ (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) | ~ (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) | ~ (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) | ~ (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) | ~ (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(subsumption_resolution, [], [f869, f2749])).
fof(f869, plain, (~ (e24 = h2(e11)) | sP55 | sP54 | sP53 | sP52 | ~ (h2(op1(e14, e14)) = op2(h2(e14), h2(e14))) | ~ (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) | ~ (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) | ~ (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) | ~ (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) | ~ (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(cnf_transformation, [], [f94])).
fof(f3367, plain, (~ spl68_434 | ~ spl68_435 | ~ spl68_436 | ~ spl68_437 | ~ spl68_438 | ~ spl68_439 | ~ spl68_440 | ~ spl68_441 | ~ spl68_442 | ~ spl68_443 | ~ spl68_444 | ~ spl68_445 | ~ spl68_446 | ~ spl68_447 | ~ spl68_448 | ~ spl68_449 | spl68_343 | spl68_338 | spl68_333 | ~ spl68_450 | ~ spl68_451 | ~ spl68_452 | ~ spl68_453 | ~ spl68_454 | ~ spl68_455 | ~ spl68_456 | ~ spl68_457 | ~ spl68_458 | ~ spl68_459), inference(avatar_split_clause, [], [f3262, f3364, f3360, f3356, f3352, f3348, f3344, f3340, f3336, f3332, f3328, f2625, f2651, f2676, f3324, f3320, f3316, f3312, f3308, f3304, f3300, f3296, f3292, f3288, f3284, f3280, f3276, f3272, f3268, f3264])).
fof(f2676, plain, (spl68_343 <=> sP56), introduced(avatar_definition, [new_symbols(naming, [spl68_343])])).
fof(f2651, plain, (spl68_338 <=> sP57), introduced(avatar_definition, [new_symbols(naming, [spl68_338])])).
fof(f2625, plain, (spl68_333 <=> sP59), introduced(avatar_definition, [new_symbols(naming, [spl68_333])])).
fof(f3262, plain, (~ (h3(op1(e10, e14)) = op2(h3(e10), e22)) | ~ (h3(e10) = h3(op1(e11, e14))) | ~ (h3(op1(e12, e14)) = op2(h3(e12), e22)) | ~ (h3(e12) = h3(op1(e13, e14))) | ~ (h3(op1(e14, e10)) = op2(e22, h3(e10))) | ~ (h3(e13) = h3(op1(e14, e11))) | ~ (h3(op1(e14, e12)) = op2(e22, h3(e12))) | ~ (h3(op1(e14, e13)) = op2(e22, h3(e13))) | ~ (h3(e11) = h3(op1(e14, e14))) | ~ (e24 = h3(e13)) | sP59 | sP57 | sP56 | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f3261, f748])).
fof(f3261, plain, (~ (h3(e10) = h3(op1(e11, e14))) | ~ (h3(op1(e12, e14)) = op2(h3(e12), e22)) | ~ (h3(e12) = h3(op1(e13, e14))) | ~ (h3(op1(e14, e10)) = op2(e22, h3(e10))) | ~ (h3(e13) = h3(op1(e14, e11))) | ~ (h3(op1(e14, e12)) = op2(e22, h3(e12))) | ~ (h3(op1(e14, e13)) = op2(e22, h3(e13))) | ~ (h3(e11) = h3(op1(e14, e14))) | ~ (e24 = h3(e13)) | sP59 | sP57 | sP56 | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f3260, f2460])).
fof(f3260, plain, (~ (h3(op1(e11, e14)) = op2(h3(e11), e22)) | ~ (h3(op1(e12, e14)) = op2(h3(e12), e22)) | ~ (h3(e12) = h3(op1(e13, e14))) | ~ (h3(op1(e14, e10)) = op2(e22, h3(e10))) | ~ (h3(e13) = h3(op1(e14, e11))) | ~ (h3(op1(e14, e12)) = op2(e22, h3(e12))) | ~ (h3(op1(e14, e13)) = op2(e22, h3(e13))) | ~ (h3(e11) = h3(op1(e14, e14))) | ~ (e24 = h3(e13)) | sP59 | sP57 | sP56 | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f3259, f748])).
fof(f3259, plain, (~ (h3(op1(e12, e14)) = op2(h3(e12), e22)) | ~ (h3(e12) = h3(op1(e13, e14))) | ~ (h3(op1(e14, e10)) = op2(e22, h3(e10))) | ~ (h3(e13) = h3(op1(e14, e11))) | ~ (h3(op1(e14, e12)) = op2(e22, h3(e12))) | ~ (h3(op1(e14, e13)) = op2(e22, h3(e13))) | ~ (h3(e11) = h3(op1(e14, e14))) | ~ (e24 = h3(e13)) | sP59 | sP57 | sP56 | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f3258, f748])).
fof(f3258, plain, (~ (h3(e12) = h3(op1(e13, e14))) | ~ (h3(op1(e14, e10)) = op2(e22, h3(e10))) | ~ (h3(e13) = h3(op1(e14, e11))) | ~ (h3(op1(e14, e12)) = op2(e22, h3(e12))) | ~ (h3(op1(e14, e13)) = op2(e22, h3(e13))) | ~ (h3(e11) = h3(op1(e14, e14))) | ~ (e24 = h3(e13)) | sP59 | sP57 | sP56 | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f3257, f2450])).
fof(f3257, plain, (~ (h3(op1(e13, e14)) = op2(h3(e13), e22)) | ~ (h3(op1(e14, e10)) = op2(e22, h3(e10))) | ~ (h3(e13) = h3(op1(e14, e11))) | ~ (h3(op1(e14, e12)) = op2(e22, h3(e12))) | ~ (h3(op1(e14, e13)) = op2(e22, h3(e13))) | ~ (h3(e11) = h3(op1(e14, e14))) | ~ (e24 = h3(e13)) | sP59 | sP57 | sP56 | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f3256, f748])).
fof(f3256, plain, (~ (h3(op1(e14, e10)) = op2(e22, h3(e10))) | ~ (h3(e13) = h3(op1(e14, e11))) | ~ (h3(op1(e14, e12)) = op2(e22, h3(e12))) | ~ (h3(op1(e14, e13)) = op2(e22, h3(e13))) | ~ (h3(e11) = h3(op1(e14, e14))) | ~ (e24 = h3(e13)) | sP59 | sP57 | sP56 | ~ (h3(op1(e13, e14)) = op2(h3(e13), h3(e14))) | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f3255, f748])).
fof(f3255, plain, (~ (h3(e13) = h3(op1(e14, e11))) | ~ (h3(op1(e14, e12)) = op2(e22, h3(e12))) | ~ (h3(op1(e14, e13)) = op2(e22, h3(e13))) | ~ (h3(e11) = h3(op1(e14, e14))) | ~ (e24 = h3(e13)) | sP59 | sP57 | sP56 | ~ (h3(op1(e14, e10)) = op2(h3(e14), h3(e10))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), h3(e14))) | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f3254, f2459])).
fof(f3254, plain, (~ (h3(op1(e14, e11)) = op2(e22, h3(e11))) | ~ (h3(op1(e14, e12)) = op2(e22, h3(e12))) | ~ (h3(op1(e14, e13)) = op2(e22, h3(e13))) | ~ (h3(e11) = h3(op1(e14, e14))) | ~ (e24 = h3(e13)) | sP59 | sP57 | sP56 | ~ (h3(op1(e14, e10)) = op2(h3(e14), h3(e10))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), h3(e14))) | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f3253, f748])).
fof(f3253, plain, (~ (h3(op1(e14, e12)) = op2(e22, h3(e12))) | ~ (h3(op1(e14, e13)) = op2(e22, h3(e13))) | ~ (h3(e11) = h3(op1(e14, e14))) | ~ (e24 = h3(e13)) | sP59 | sP57 | sP56 | ~ (h3(op1(e14, e11)) = op2(h3(e14), h3(e11))) | ~ (h3(op1(e14, e10)) = op2(h3(e14), h3(e10))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), h3(e14))) | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f3252, f748])).
fof(f3252, plain, (~ (h3(op1(e14, e13)) = op2(e22, h3(e13))) | ~ (h3(e11) = h3(op1(e14, e14))) | ~ (e24 = h3(e13)) | sP59 | sP57 | sP56 | ~ (h3(op1(e14, e12)) = op2(h3(e14), h3(e12))) | ~ (h3(op1(e14, e11)) = op2(h3(e14), h3(e11))) | ~ (h3(op1(e14, e10)) = op2(h3(e14), h3(e10))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), h3(e14))) | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f3251, f748])).
fof(f3251, plain, (~ (h3(e11) = h3(op1(e14, e14))) | ~ (e24 = h3(e13)) | sP59 | sP57 | sP56 | ~ (h3(op1(e14, e13)) = op2(h3(e14), h3(e13))) | ~ (h3(op1(e14, e12)) = op2(h3(e14), h3(e12))) | ~ (h3(op1(e14, e11)) = op2(h3(e14), h3(e11))) | ~ (h3(op1(e14, e10)) = op2(h3(e14), h3(e10))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), h3(e14))) | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f3250, f750])).
fof(f3250, plain, (~ (op2(e22, e22) = h3(op1(e14, e14))) | ~ (e24 = h3(e13)) | sP59 | sP57 | sP56 | ~ (h3(op1(e14, e13)) = op2(h3(e14), h3(e13))) | ~ (h3(op1(e14, e12)) = op2(h3(e14), h3(e12))) | ~ (h3(op1(e14, e11)) = op2(h3(e14), h3(e11))) | ~ (h3(op1(e14, e10)) = op2(h3(e14), h3(e10))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), h3(e14))) | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f3249, f748])).
fof(f3249, plain, (~ (e24 = h3(e13)) | sP59 | sP57 | sP56 | ~ (h3(op1(e14, e14)) = op2(h3(e14), h3(e14))) | ~ (h3(op1(e14, e13)) = op2(h3(e14), h3(e13))) | ~ (h3(op1(e14, e12)) = op2(h3(e14), h3(e12))) | ~ (h3(op1(e14, e11)) = op2(h3(e14), h3(e11))) | ~ (h3(op1(e14, e10)) = op2(h3(e14), h3(e10))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), h3(e14))) | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(subsumption_resolution, [], [f876, f2648])).
fof(f2648, plain, ~ sP58, inference(subsumption_resolution, [], [f812, f748])).
fof(f812, plain, (~ (e22 = h3(e14)) | ~ sP58), inference(cnf_transformation, [], [f152])).
fof(f152, plain, ((~ (e22 = h3(e14)) & ~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10))) | ~ sP58), inference(nnf_transformation, [], [f84])).
fof(f876, plain, (~ (e24 = h3(e13)) | sP59 | sP58 | sP57 | sP56 | ~ (h3(op1(e14, e14)) = op2(h3(e14), h3(e14))) | ~ (h3(op1(e14, e13)) = op2(h3(e14), h3(e13))) | ~ (h3(op1(e14, e12)) = op2(h3(e14), h3(e12))) | ~ (h3(op1(e14, e11)) = op2(h3(e14), h3(e11))) | ~ (h3(op1(e14, e10)) = op2(h3(e14), h3(e10))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), h3(e14))) | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(cnf_transformation, [], [f94])).
fof(f3208, plain, (~ spl68_405 | ~ spl68_406 | ~ spl68_407 | ~ spl68_408 | ~ spl68_409 | ~ spl68_410 | ~ spl68_411 | ~ spl68_412 | ~ spl68_413 | ~ spl68_414 | ~ spl68_415 | ~ spl68_416 | ~ spl68_417 | ~ spl68_418 | ~ spl68_419 | ~ spl68_420 | spl68_328 | spl68_323 | spl68_318 | ~ spl68_431 | ~ spl68_422 | ~ spl68_423 | ~ spl68_424 | ~ spl68_425 | ~ spl68_426 | ~ spl68_427 | ~ spl68_428 | ~ spl68_429 | ~ spl68_430), inference(avatar_split_clause, [], [f3203, f3186, f3182, f3178, f3174, f3170, f3166, f3162, f3158, f3154, f3205, f2550, f2575, f2600, f3146, f3142, f3138, f3134, f3130, f3126, f3122, f3118, f3114, f3110, f3106, f3102, f3098, f3094, f3090, f3086])).
fof(f2600, plain, (spl68_328 <=> sP60), introduced(avatar_definition, [new_symbols(naming, [spl68_328])])).
fof(f2575, plain, (spl68_323 <=> sP61), introduced(avatar_definition, [new_symbols(naming, [spl68_323])])).
fof(f2550, plain, (spl68_318 <=> sP62), introduced(avatar_definition, [new_symbols(naming, [spl68_318])])).
fof(f3203, plain, (~ (h4(op1(e10, e14)) = op2(h4(e10), e23)) | ~ (h4(e10) = h4(op1(e11, e14))) | ~ (h4(op1(e12, e14)) = op2(h4(e12), e23)) | ~ (h4(e12) = h4(op1(e13, e14))) | ~ (h4(op1(e14, e10)) = op2(e23, h4(e10))) | ~ (h4(e13) = h4(op1(e14, e11))) | ~ (h4(op1(e14, e12)) = op2(e23, h4(e12))) | ~ (h4(op1(e14, e13)) = op2(e23, h4(e13))) | ~ (h4(e11) = h4(op1(e14, e14))) | ~ (e24 = h4(e12)) | sP62 | sP61 | sP60 | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f3202, f753])).
fof(f3202, plain, (~ (h4(e10) = h4(op1(e11, e14))) | ~ (h4(op1(e12, e14)) = op2(h4(e12), e23)) | ~ (h4(e12) = h4(op1(e13, e14))) | ~ (h4(op1(e14, e10)) = op2(e23, h4(e10))) | ~ (h4(e13) = h4(op1(e14, e11))) | ~ (h4(op1(e14, e12)) = op2(e23, h4(e12))) | ~ (h4(op1(e14, e13)) = op2(e23, h4(e13))) | ~ (h4(e11) = h4(op1(e14, e14))) | ~ (e24 = h4(e12)) | sP62 | sP61 | sP60 | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e14)) = op2(h4(e10), h4(e14))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f3201, f2471])).
fof(f3201, plain, (~ (h4(op1(e11, e14)) = op2(h4(e11), e23)) | ~ (h4(op1(e12, e14)) = op2(h4(e12), e23)) | ~ (h4(e12) = h4(op1(e13, e14))) | ~ (h4(op1(e14, e10)) = op2(e23, h4(e10))) | ~ (h4(e13) = h4(op1(e14, e11))) | ~ (h4(op1(e14, e12)) = op2(e23, h4(e12))) | ~ (h4(op1(e14, e13)) = op2(e23, h4(e13))) | ~ (h4(e11) = h4(op1(e14, e14))) | ~ (e24 = h4(e12)) | sP62 | sP61 | sP60 | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e14)) = op2(h4(e10), h4(e14))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f3200, f753])).
fof(f3200, plain, (~ (h4(op1(e12, e14)) = op2(h4(e12), e23)) | ~ (h4(e12) = h4(op1(e13, e14))) | ~ (h4(op1(e14, e10)) = op2(e23, h4(e10))) | ~ (h4(e13) = h4(op1(e14, e11))) | ~ (h4(op1(e14, e12)) = op2(e23, h4(e12))) | ~ (h4(op1(e14, e13)) = op2(e23, h4(e13))) | ~ (h4(e11) = h4(op1(e14, e14))) | ~ (e24 = h4(e12)) | sP62 | sP61 | sP60 | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e14)) = op2(h4(e11), h4(e14))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e14)) = op2(h4(e10), h4(e14))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f3199, f753])).
fof(f3199, plain, (~ (h4(e12) = h4(op1(e13, e14))) | ~ (h4(op1(e14, e10)) = op2(e23, h4(e10))) | ~ (h4(e13) = h4(op1(e14, e11))) | ~ (h4(op1(e14, e12)) = op2(e23, h4(e12))) | ~ (h4(op1(e14, e13)) = op2(e23, h4(e13))) | ~ (h4(e11) = h4(op1(e14, e14))) | ~ (e24 = h4(e12)) | sP62 | sP61 | sP60 | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e14)) = op2(h4(e12), h4(e14))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e14)) = op2(h4(e11), h4(e14))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e14)) = op2(h4(e10), h4(e14))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f3198, f2461])).
fof(f3198, plain, (~ (h4(op1(e13, e14)) = op2(h4(e13), e23)) | ~ (h4(op1(e14, e10)) = op2(e23, h4(e10))) | ~ (h4(e13) = h4(op1(e14, e11))) | ~ (h4(op1(e14, e12)) = op2(e23, h4(e12))) | ~ (h4(op1(e14, e13)) = op2(e23, h4(e13))) | ~ (h4(e11) = h4(op1(e14, e14))) | ~ (e24 = h4(e12)) | sP62 | sP61 | sP60 | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e14)) = op2(h4(e12), h4(e14))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e14)) = op2(h4(e11), h4(e14))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e14)) = op2(h4(e10), h4(e14))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f3197, f753])).
fof(f3197, plain, (~ (h4(op1(e14, e10)) = op2(e23, h4(e10))) | ~ (h4(e13) = h4(op1(e14, e11))) | ~ (h4(op1(e14, e12)) = op2(e23, h4(e12))) | ~ (h4(op1(e14, e13)) = op2(e23, h4(e13))) | ~ (h4(e11) = h4(op1(e14, e14))) | ~ (e24 = h4(e12)) | sP62 | sP61 | sP60 | ~ (h4(op1(e13, e14)) = op2(h4(e13), h4(e14))) | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e14)) = op2(h4(e12), h4(e14))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e14)) = op2(h4(e11), h4(e14))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e14)) = op2(h4(e10), h4(e14))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f3196, f753])).
fof(f3196, plain, (~ (h4(e13) = h4(op1(e14, e11))) | ~ (h4(op1(e14, e12)) = op2(e23, h4(e12))) | ~ (h4(op1(e14, e13)) = op2(e23, h4(e13))) | ~ (h4(e11) = h4(op1(e14, e14))) | ~ (e24 = h4(e12)) | sP62 | sP61 | sP60 | ~ (h4(op1(e14, e10)) = op2(h4(e14), h4(e10))) | ~ (h4(op1(e13, e14)) = op2(h4(e13), h4(e14))) | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e14)) = op2(h4(e12), h4(e14))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e14)) = op2(h4(e11), h4(e14))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e14)) = op2(h4(e10), h4(e14))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f3195, f2470])).
fof(f3195, plain, (~ (h4(op1(e14, e11)) = op2(e23, h4(e11))) | ~ (h4(op1(e14, e12)) = op2(e23, h4(e12))) | ~ (h4(op1(e14, e13)) = op2(e23, h4(e13))) | ~ (h4(e11) = h4(op1(e14, e14))) | ~ (e24 = h4(e12)) | sP62 | sP61 | sP60 | ~ (h4(op1(e14, e10)) = op2(h4(e14), h4(e10))) | ~ (h4(op1(e13, e14)) = op2(h4(e13), h4(e14))) | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e14)) = op2(h4(e12), h4(e14))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e14)) = op2(h4(e11), h4(e14))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e14)) = op2(h4(e10), h4(e14))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f3194, f753])).
fof(f3194, plain, (~ (h4(op1(e14, e12)) = op2(e23, h4(e12))) | ~ (h4(op1(e14, e13)) = op2(e23, h4(e13))) | ~ (h4(e11) = h4(op1(e14, e14))) | ~ (e24 = h4(e12)) | sP62 | sP61 | sP60 | ~ (h4(op1(e14, e11)) = op2(h4(e14), h4(e11))) | ~ (h4(op1(e14, e10)) = op2(h4(e14), h4(e10))) | ~ (h4(op1(e13, e14)) = op2(h4(e13), h4(e14))) | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e14)) = op2(h4(e12), h4(e14))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e14)) = op2(h4(e11), h4(e14))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e14)) = op2(h4(e10), h4(e14))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f3193, f753])).
fof(f3193, plain, (~ (h4(op1(e14, e13)) = op2(e23, h4(e13))) | ~ (h4(e11) = h4(op1(e14, e14))) | ~ (e24 = h4(e12)) | sP62 | sP61 | sP60 | ~ (h4(op1(e14, e12)) = op2(h4(e14), h4(e12))) | ~ (h4(op1(e14, e11)) = op2(h4(e14), h4(e11))) | ~ (h4(op1(e14, e10)) = op2(h4(e14), h4(e10))) | ~ (h4(op1(e13, e14)) = op2(h4(e13), h4(e14))) | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e14)) = op2(h4(e12), h4(e14))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e14)) = op2(h4(e11), h4(e14))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e14)) = op2(h4(e10), h4(e14))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f3192, f753])).
fof(f3192, plain, (~ (h4(e11) = h4(op1(e14, e14))) | ~ (e24 = h4(e12)) | sP62 | sP61 | sP60 | ~ (h4(op1(e14, e13)) = op2(h4(e14), h4(e13))) | ~ (h4(op1(e14, e12)) = op2(h4(e14), h4(e12))) | ~ (h4(op1(e14, e11)) = op2(h4(e14), h4(e11))) | ~ (h4(op1(e14, e10)) = op2(h4(e14), h4(e10))) | ~ (h4(op1(e13, e14)) = op2(h4(e13), h4(e14))) | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e14)) = op2(h4(e12), h4(e14))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e14)) = op2(h4(e11), h4(e14))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e14)) = op2(h4(e10), h4(e14))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f3191, f755])).
fof(f3191, plain, (~ (op2(e23, e23) = h4(op1(e14, e14))) | ~ (e24 = h4(e12)) | sP62 | sP61 | sP60 | ~ (h4(op1(e14, e13)) = op2(h4(e14), h4(e13))) | ~ (h4(op1(e14, e12)) = op2(h4(e14), h4(e12))) | ~ (h4(op1(e14, e11)) = op2(h4(e14), h4(e11))) | ~ (h4(op1(e14, e10)) = op2(h4(e14), h4(e10))) | ~ (h4(op1(e13, e14)) = op2(h4(e13), h4(e14))) | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e14)) = op2(h4(e12), h4(e14))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e14)) = op2(h4(e11), h4(e14))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e14)) = op2(h4(e10), h4(e14))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f3190, f753])).
fof(f3190, plain, (~ (e24 = h4(e12)) | sP62 | sP61 | sP60 | ~ (h4(op1(e14, e14)) = op2(h4(e14), h4(e14))) | ~ (h4(op1(e14, e13)) = op2(h4(e14), h4(e13))) | ~ (h4(op1(e14, e12)) = op2(h4(e14), h4(e12))) | ~ (h4(op1(e14, e11)) = op2(h4(e14), h4(e11))) | ~ (h4(op1(e14, e10)) = op2(h4(e14), h4(e10))) | ~ (h4(op1(e13, e14)) = op2(h4(e13), h4(e14))) | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e14)) = op2(h4(e12), h4(e14))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e14)) = op2(h4(e11), h4(e14))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e14)) = op2(h4(e10), h4(e14))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(subsumption_resolution, [], [f880, f2547])).
fof(f2547, plain, ~ sP63, inference(subsumption_resolution, [], [f787, f753])).
fof(f787, plain, (~ (e23 = h4(e14)) | ~ sP63), inference(cnf_transformation, [], [f147])).
fof(f147, plain, ((~ (e23 = h4(e14)) & ~ (e23 = h4(e13)) & ~ (e23 = h4(e12)) & ~ (e23 = h4(e11)) & ~ (e23 = h4(e10))) | ~ sP63), inference(nnf_transformation, [], [f89])).
fof(f880, plain, (~ (e24 = h4(e12)) | sP63 | sP62 | sP61 | sP60 | ~ (h4(op1(e14, e14)) = op2(h4(e14), h4(e14))) | ~ (h4(op1(e14, e13)) = op2(h4(e14), h4(e13))) | ~ (h4(op1(e14, e12)) = op2(h4(e14), h4(e12))) | ~ (h4(op1(e14, e11)) = op2(h4(e14), h4(e11))) | ~ (h4(op1(e14, e10)) = op2(h4(e14), h4(e10))) | ~ (h4(op1(e13, e14)) = op2(h4(e13), h4(e14))) | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e14)) = op2(h4(e12), h4(e14))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e14)) = op2(h4(e11), h4(e14))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e14)) = op2(h4(e10), h4(e14))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(cnf_transformation, [], [f94])).
fof(f2986, plain, (~ spl68_378 | ~ spl68_379 | ~ spl68_380 | ~ spl68_381 | spl68_315 | spl68_312 | spl68_309 | ~ spl68_382 | ~ spl68_383 | ~ spl68_384 | ~ spl68_385 | ~ spl68_386 | ~ spl68_387 | ~ spl68_388 | ~ spl68_389 | ~ spl68_390 | ~ spl68_391 | ~ spl68_392 | ~ spl68_393 | ~ spl68_394 | ~ spl68_395 | ~ spl68_396 | ~ spl68_397 | ~ spl68_398 | ~ spl68_399 | ~ spl68_400 | ~ spl68_401 | ~ spl68_402), inference(avatar_split_clause, [], [f2885, f2983, f2979, f2975, f2971, f2967, f2963, f2959, f2955, f2951, f2947, f2943, f2939, f2935, f2931, f2927, f2923, f2919, f2915, f2911, f2907, f2903, f2498, f2515, f2532, f2899, f2895, f2891, f2887])).
fof(f2532, plain, (spl68_315 <=> sP64), introduced(avatar_definition, [new_symbols(naming, [spl68_315])])).
fof(f2515, plain, (spl68_312 <=> sP65), introduced(avatar_definition, [new_symbols(naming, [spl68_312])])).
fof(f2498, plain, (spl68_309 <=> sP66), introduced(avatar_definition, [new_symbols(naming, [spl68_309])])).
fof(f2885, plain, (~ (h1(e11) = h5(op1(e10, e10))) | ~ (h5(op1(e10, e11)) = op2(e20, h5(e11))) | ~ (h5(op1(e10, e12)) = op2(e20, h5(e12))) | ~ (op2(e20, e23) = h5(op1(e10, e13))) | ~ (op2(e20, e24) = h5(op1(e10, e14))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), e20)) | ~ (h5(op1(e11, e13)) = op2(h5(e11), e23)) | ~ (e20 = h5(op1(e11, e14))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), e20)) | ~ (h5(op1(e12, e13)) = op2(h5(e12), e23)) | ~ (h5(op1(e12, e14)) = op2(h5(e12), e24)) | ~ (op2(e23, e20) = h5(op1(e13, e10))) | ~ (h5(op1(e13, e11)) = op2(e23, h5(e11))) | ~ (h5(op1(e13, e12)) = op2(e23, h5(e12))) | ~ (h4(e11) = h5(op1(e13, e13))) | ~ (h5(e12) = h5(op1(e13, e14))) | ~ (op2(e24, e20) = h5(op1(e14, e10))) | ~ (e23 = h5(op1(e14, e11))) | ~ (h5(op1(e14, e12)) = op2(e24, h5(e12))) | ~ (op2(e24, e23) = h5(op1(e14, e13))) | ~ (h5(e11) = h5(op1(e14, e14))) | sP66 | sP65 | sP64 | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11)))), inference(forward_demodulation, [], [f2884, f740])).
fof(f2884, plain, (~ (op2(e20, e20) = h5(op1(e10, e10))) | ~ (h5(op1(e10, e11)) = op2(e20, h5(e11))) | ~ (h5(op1(e10, e12)) = op2(e20, h5(e12))) | ~ (op2(e20, e23) = h5(op1(e10, e13))) | ~ (op2(e20, e24) = h5(op1(e10, e14))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), e20)) | ~ (h5(op1(e11, e13)) = op2(h5(e11), e23)) | ~ (e20 = h5(op1(e11, e14))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), e20)) | ~ (h5(op1(e12, e13)) = op2(h5(e12), e23)) | ~ (h5(op1(e12, e14)) = op2(h5(e12), e24)) | ~ (op2(e23, e20) = h5(op1(e13, e10))) | ~ (h5(op1(e13, e11)) = op2(e23, h5(e11))) | ~ (h5(op1(e13, e12)) = op2(e23, h5(e12))) | ~ (h4(e11) = h5(op1(e13, e13))) | ~ (h5(e12) = h5(op1(e13, e14))) | ~ (op2(e24, e20) = h5(op1(e14, e10))) | ~ (e23 = h5(op1(e14, e11))) | ~ (h5(op1(e14, e12)) = op2(e24, h5(e12))) | ~ (op2(e24, e23) = h5(op1(e14, e13))) | ~ (h5(e11) = h5(op1(e14, e14))) | sP66 | sP65 | sP64 | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11)))), inference(forward_demodulation, [], [f2883, f2492])).
fof(f2883, plain, (~ (h5(op1(e10, e11)) = op2(e20, h5(e11))) | ~ (h5(op1(e10, e12)) = op2(e20, h5(e12))) | ~ (op2(e20, e23) = h5(op1(e10, e13))) | ~ (op2(e20, e24) = h5(op1(e10, e14))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), e20)) | ~ (h5(op1(e11, e13)) = op2(h5(e11), e23)) | ~ (e20 = h5(op1(e11, e14))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), e20)) | ~ (h5(op1(e12, e13)) = op2(h5(e12), e23)) | ~ (h5(op1(e12, e14)) = op2(h5(e12), e24)) | ~ (op2(e23, e20) = h5(op1(e13, e10))) | ~ (h5(op1(e13, e11)) = op2(e23, h5(e11))) | ~ (h5(op1(e13, e12)) = op2(e23, h5(e12))) | ~ (h4(e11) = h5(op1(e13, e13))) | ~ (h5(e12) = h5(op1(e13, e14))) | ~ (op2(e24, e20) = h5(op1(e14, e10))) | ~ (e23 = h5(op1(e14, e11))) | ~ (h5(op1(e14, e12)) = op2(e24, h5(e12))) | ~ (op2(e24, e23) = h5(op1(e14, e13))) | ~ (h5(e11) = h5(op1(e14, e14))) | sP66 | sP65 | sP64 | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f2882, f2492])).
fof(f2882, plain, (~ (h5(op1(e10, e12)) = op2(e20, h5(e12))) | ~ (op2(e20, e23) = h5(op1(e10, e13))) | ~ (op2(e20, e24) = h5(op1(e10, e14))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), e20)) | ~ (h5(op1(e11, e13)) = op2(h5(e11), e23)) | ~ (e20 = h5(op1(e11, e14))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), e20)) | ~ (h5(op1(e12, e13)) = op2(h5(e12), e23)) | ~ (h5(op1(e12, e14)) = op2(h5(e12), e24)) | ~ (op2(e23, e20) = h5(op1(e13, e10))) | ~ (h5(op1(e13, e11)) = op2(e23, h5(e11))) | ~ (h5(op1(e13, e12)) = op2(e23, h5(e12))) | ~ (h4(e11) = h5(op1(e13, e13))) | ~ (h5(e12) = h5(op1(e13, e14))) | ~ (op2(e24, e20) = h5(op1(e14, e10))) | ~ (e23 = h5(op1(e14, e11))) | ~ (h5(op1(e14, e12)) = op2(e24, h5(e12))) | ~ (op2(e24, e23) = h5(op1(e14, e13))) | ~ (h5(e11) = h5(op1(e14, e14))) | sP66 | sP65 | sP64 | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f2881, f2492])).
fof(f2881, plain, (~ (op2(e20, e23) = h5(op1(e10, e13))) | ~ (op2(e20, e24) = h5(op1(e10, e14))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), e20)) | ~ (h5(op1(e11, e13)) = op2(h5(e11), e23)) | ~ (e20 = h5(op1(e11, e14))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), e20)) | ~ (h5(op1(e12, e13)) = op2(h5(e12), e23)) | ~ (h5(op1(e12, e14)) = op2(h5(e12), e24)) | ~ (op2(e23, e20) = h5(op1(e13, e10))) | ~ (h5(op1(e13, e11)) = op2(e23, h5(e11))) | ~ (h5(op1(e13, e12)) = op2(e23, h5(e12))) | ~ (h4(e11) = h5(op1(e13, e13))) | ~ (h5(e12) = h5(op1(e13, e14))) | ~ (op2(e24, e20) = h5(op1(e14, e10))) | ~ (e23 = h5(op1(e14, e11))) | ~ (h5(op1(e14, e12)) = op2(e24, h5(e12))) | ~ (op2(e24, e23) = h5(op1(e14, e13))) | ~ (h5(e11) = h5(op1(e14, e14))) | sP66 | sP65 | sP64 | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f2880, f2492])).
fof(f2880, plain, (~ (h5(op1(e10, e13)) = op2(h5(e10), e23)) | ~ (op2(e20, e24) = h5(op1(e10, e14))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), e20)) | ~ (h5(op1(e11, e13)) = op2(h5(e11), e23)) | ~ (e20 = h5(op1(e11, e14))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), e20)) | ~ (h5(op1(e12, e13)) = op2(h5(e12), e23)) | ~ (h5(op1(e12, e14)) = op2(h5(e12), e24)) | ~ (op2(e23, e20) = h5(op1(e13, e10))) | ~ (h5(op1(e13, e11)) = op2(e23, h5(e11))) | ~ (h5(op1(e13, e12)) = op2(e23, h5(e12))) | ~ (h4(e11) = h5(op1(e13, e13))) | ~ (h5(e12) = h5(op1(e13, e14))) | ~ (op2(e24, e20) = h5(op1(e14, e10))) | ~ (e23 = h5(op1(e14, e11))) | ~ (h5(op1(e14, e12)) = op2(e24, h5(e12))) | ~ (op2(e24, e23) = h5(op1(e14, e13))) | ~ (h5(e11) = h5(op1(e14, e14))) | sP66 | sP65 | sP64 | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f2879, f2472])).
fof(f2879, plain, (~ (op2(e20, e24) = h5(op1(e10, e14))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), e20)) | ~ (h5(op1(e11, e13)) = op2(h5(e11), e23)) | ~ (e20 = h5(op1(e11, e14))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), e20)) | ~ (h5(op1(e12, e13)) = op2(h5(e12), e23)) | ~ (h5(op1(e12, e14)) = op2(h5(e12), e24)) | ~ (op2(e23, e20) = h5(op1(e13, e10))) | ~ (h5(op1(e13, e11)) = op2(e23, h5(e11))) | ~ (h5(op1(e13, e12)) = op2(e23, h5(e12))) | ~ (h4(e11) = h5(op1(e13, e13))) | ~ (h5(e12) = h5(op1(e13, e14))) | ~ (op2(e24, e20) = h5(op1(e14, e10))) | ~ (e23 = h5(op1(e14, e11))) | ~ (h5(op1(e14, e12)) = op2(e24, h5(e12))) | ~ (op2(e24, e23) = h5(op1(e14, e13))) | ~ (h5(e11) = h5(op1(e14, e14))) | sP66 | sP65 | sP64 | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f2878, f2492])).
fof(f2878, plain, (~ (h5(op1(e10, e14)) = op2(h5(e10), e24)) | ~ (h5(op1(e11, e10)) = op2(h5(e11), e20)) | ~ (h5(op1(e11, e13)) = op2(h5(e11), e23)) | ~ (e20 = h5(op1(e11, e14))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), e20)) | ~ (h5(op1(e12, e13)) = op2(h5(e12), e23)) | ~ (h5(op1(e12, e14)) = op2(h5(e12), e24)) | ~ (op2(e23, e20) = h5(op1(e13, e10))) | ~ (h5(op1(e13, e11)) = op2(e23, h5(e11))) | ~ (h5(op1(e13, e12)) = op2(e23, h5(e12))) | ~ (h4(e11) = h5(op1(e13, e13))) | ~ (h5(e12) = h5(op1(e13, e14))) | ~ (op2(e24, e20) = h5(op1(e14, e10))) | ~ (e23 = h5(op1(e14, e11))) | ~ (h5(op1(e14, e12)) = op2(e24, h5(e12))) | ~ (op2(e24, e23) = h5(op1(e14, e13))) | ~ (h5(e11) = h5(op1(e14, e14))) | sP66 | sP65 | sP64 | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f2877, f758])).
fof(f2877, plain, (~ (h5(op1(e11, e10)) = op2(h5(e11), e20)) | ~ (h5(op1(e11, e13)) = op2(h5(e11), e23)) | ~ (e20 = h5(op1(e11, e14))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), e20)) | ~ (h5(op1(e12, e13)) = op2(h5(e12), e23)) | ~ (h5(op1(e12, e14)) = op2(h5(e12), e24)) | ~ (op2(e23, e20) = h5(op1(e13, e10))) | ~ (h5(op1(e13, e11)) = op2(e23, h5(e11))) | ~ (h5(op1(e13, e12)) = op2(e23, h5(e12))) | ~ (h4(e11) = h5(op1(e13, e13))) | ~ (h5(e12) = h5(op1(e13, e14))) | ~ (op2(e24, e20) = h5(op1(e14, e10))) | ~ (e23 = h5(op1(e14, e11))) | ~ (h5(op1(e14, e12)) = op2(e24, h5(e12))) | ~ (op2(e24, e23) = h5(op1(e14, e13))) | ~ (h5(e11) = h5(op1(e14, e14))) | sP66 | sP65 | sP64 | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f2876, f2492])).
fof(f2876, plain, (~ (h5(op1(e11, e13)) = op2(h5(e11), e23)) | ~ (e20 = h5(op1(e11, e14))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), e20)) | ~ (h5(op1(e12, e13)) = op2(h5(e12), e23)) | ~ (h5(op1(e12, e14)) = op2(h5(e12), e24)) | ~ (op2(e23, e20) = h5(op1(e13, e10))) | ~ (h5(op1(e13, e11)) = op2(e23, h5(e11))) | ~ (h5(op1(e13, e12)) = op2(e23, h5(e12))) | ~ (h4(e11) = h5(op1(e13, e13))) | ~ (h5(e12) = h5(op1(e13, e14))) | ~ (op2(e24, e20) = h5(op1(e14, e10))) | ~ (e23 = h5(op1(e14, e11))) | ~ (h5(op1(e14, e12)) = op2(e24, h5(e12))) | ~ (op2(e24, e23) = h5(op1(e14, e13))) | ~ (h5(e11) = h5(op1(e14, e14))) | sP66 | sP65 | sP64 | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f2875, f2472])).
fof(f2875, plain, (~ (e20 = h5(op1(e11, e14))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), e20)) | ~ (h5(op1(e12, e13)) = op2(h5(e12), e23)) | ~ (h5(op1(e12, e14)) = op2(h5(e12), e24)) | ~ (op2(e23, e20) = h5(op1(e13, e10))) | ~ (h5(op1(e13, e11)) = op2(e23, h5(e11))) | ~ (h5(op1(e13, e12)) = op2(e23, h5(e12))) | ~ (h4(e11) = h5(op1(e13, e13))) | ~ (h5(e12) = h5(op1(e13, e14))) | ~ (op2(e24, e20) = h5(op1(e14, e10))) | ~ (e23 = h5(op1(e14, e11))) | ~ (h5(op1(e14, e12)) = op2(e24, h5(e12))) | ~ (op2(e24, e23) = h5(op1(e14, e13))) | ~ (h5(e11) = h5(op1(e14, e14))) | sP66 | sP65 | sP64 | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f2874, f2488])).
fof(f2874, plain, (~ (h5(op1(e11, e14)) = op2(h5(e11), e24)) | ~ (h5(op1(e12, e10)) = op2(h5(e12), e20)) | ~ (h5(op1(e12, e13)) = op2(h5(e12), e23)) | ~ (h5(op1(e12, e14)) = op2(h5(e12), e24)) | ~ (op2(e23, e20) = h5(op1(e13, e10))) | ~ (h5(op1(e13, e11)) = op2(e23, h5(e11))) | ~ (h5(op1(e13, e12)) = op2(e23, h5(e12))) | ~ (h4(e11) = h5(op1(e13, e13))) | ~ (h5(e12) = h5(op1(e13, e14))) | ~ (op2(e24, e20) = h5(op1(e14, e10))) | ~ (e23 = h5(op1(e14, e11))) | ~ (h5(op1(e14, e12)) = op2(e24, h5(e12))) | ~ (op2(e24, e23) = h5(op1(e14, e13))) | ~ (h5(e11) = h5(op1(e14, e14))) | sP66 | sP65 | sP64 | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f2873, f758])).
fof(f2873, plain, (~ (h5(op1(e12, e10)) = op2(h5(e12), e20)) | ~ (h5(op1(e12, e13)) = op2(h5(e12), e23)) | ~ (h5(op1(e12, e14)) = op2(h5(e12), e24)) | ~ (op2(e23, e20) = h5(op1(e13, e10))) | ~ (h5(op1(e13, e11)) = op2(e23, h5(e11))) | ~ (h5(op1(e13, e12)) = op2(e23, h5(e12))) | ~ (h4(e11) = h5(op1(e13, e13))) | ~ (h5(e12) = h5(op1(e13, e14))) | ~ (op2(e24, e20) = h5(op1(e14, e10))) | ~ (e23 = h5(op1(e14, e11))) | ~ (h5(op1(e14, e12)) = op2(e24, h5(e12))) | ~ (op2(e24, e23) = h5(op1(e14, e13))) | ~ (h5(e11) = h5(op1(e14, e14))) | sP66 | sP65 | sP64 | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f2872, f2492])).
fof(f2872, plain, (~ (h5(op1(e12, e13)) = op2(h5(e12), e23)) | ~ (h5(op1(e12, e14)) = op2(h5(e12), e24)) | ~ (op2(e23, e20) = h5(op1(e13, e10))) | ~ (h5(op1(e13, e11)) = op2(e23, h5(e11))) | ~ (h5(op1(e13, e12)) = op2(e23, h5(e12))) | ~ (h4(e11) = h5(op1(e13, e13))) | ~ (h5(e12) = h5(op1(e13, e14))) | ~ (op2(e24, e20) = h5(op1(e14, e10))) | ~ (e23 = h5(op1(e14, e11))) | ~ (h5(op1(e14, e12)) = op2(e24, h5(e12))) | ~ (op2(e24, e23) = h5(op1(e14, e13))) | ~ (h5(e11) = h5(op1(e14, e14))) | sP66 | sP65 | sP64 | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f2871, f2472])).
fof(f2871, plain, (~ (h5(op1(e12, e14)) = op2(h5(e12), e24)) | ~ (op2(e23, e20) = h5(op1(e13, e10))) | ~ (h5(op1(e13, e11)) = op2(e23, h5(e11))) | ~ (h5(op1(e13, e12)) = op2(e23, h5(e12))) | ~ (h4(e11) = h5(op1(e13, e13))) | ~ (h5(e12) = h5(op1(e13, e14))) | ~ (op2(e24, e20) = h5(op1(e14, e10))) | ~ (e23 = h5(op1(e14, e11))) | ~ (h5(op1(e14, e12)) = op2(e24, h5(e12))) | ~ (op2(e24, e23) = h5(op1(e14, e13))) | ~ (h5(e11) = h5(op1(e14, e14))) | sP66 | sP65 | sP64 | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f2870, f758])).
fof(f2870, plain, (~ (op2(e23, e20) = h5(op1(e13, e10))) | ~ (h5(op1(e13, e11)) = op2(e23, h5(e11))) | ~ (h5(op1(e13, e12)) = op2(e23, h5(e12))) | ~ (h4(e11) = h5(op1(e13, e13))) | ~ (h5(e12) = h5(op1(e13, e14))) | ~ (op2(e24, e20) = h5(op1(e14, e10))) | ~ (e23 = h5(op1(e14, e11))) | ~ (h5(op1(e14, e12)) = op2(e24, h5(e12))) | ~ (op2(e24, e23) = h5(op1(e14, e13))) | ~ (h5(e11) = h5(op1(e14, e14))) | sP66 | sP65 | sP64 | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f2869, f2472])).
fof(f2869, plain, (~ (h5(op1(e13, e10)) = op2(h5(e13), e20)) | ~ (h5(op1(e13, e11)) = op2(e23, h5(e11))) | ~ (h5(op1(e13, e12)) = op2(e23, h5(e12))) | ~ (h4(e11) = h5(op1(e13, e13))) | ~ (h5(e12) = h5(op1(e13, e14))) | ~ (op2(e24, e20) = h5(op1(e14, e10))) | ~ (e23 = h5(op1(e14, e11))) | ~ (h5(op1(e14, e12)) = op2(e24, h5(e12))) | ~ (op2(e24, e23) = h5(op1(e14, e13))) | ~ (h5(e11) = h5(op1(e14, e14))) | sP66 | sP65 | sP64 | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f2868, f2492])).
fof(f2868, plain, (~ (h5(op1(e13, e11)) = op2(e23, h5(e11))) | ~ (h5(op1(e13, e12)) = op2(e23, h5(e12))) | ~ (h4(e11) = h5(op1(e13, e13))) | ~ (h5(e12) = h5(op1(e13, e14))) | ~ (op2(e24, e20) = h5(op1(e14, e10))) | ~ (e23 = h5(op1(e14, e11))) | ~ (h5(op1(e14, e12)) = op2(e24, h5(e12))) | ~ (op2(e24, e23) = h5(op1(e14, e13))) | ~ (h5(e11) = h5(op1(e14, e14))) | sP66 | sP65 | sP64 | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f2867, f2472])).
fof(f2867, plain, (~ (h5(op1(e13, e12)) = op2(e23, h5(e12))) | ~ (h4(e11) = h5(op1(e13, e13))) | ~ (h5(e12) = h5(op1(e13, e14))) | ~ (op2(e24, e20) = h5(op1(e14, e10))) | ~ (e23 = h5(op1(e14, e11))) | ~ (h5(op1(e14, e12)) = op2(e24, h5(e12))) | ~ (op2(e24, e23) = h5(op1(e14, e13))) | ~ (h5(e11) = h5(op1(e14, e14))) | sP66 | sP65 | sP64 | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f2866, f2472])).
fof(f2866, plain, (~ (h4(e11) = h5(op1(e13, e13))) | ~ (h5(e12) = h5(op1(e13, e14))) | ~ (op2(e24, e20) = h5(op1(e14, e10))) | ~ (e23 = h5(op1(e14, e11))) | ~ (h5(op1(e14, e12)) = op2(e24, h5(e12))) | ~ (op2(e24, e23) = h5(op1(e14, e13))) | ~ (h5(e11) = h5(op1(e14, e14))) | sP66 | sP65 | sP64 | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f2865, f755])).
fof(f2865, plain, (~ (op2(e23, e23) = h5(op1(e13, e13))) | ~ (h5(e12) = h5(op1(e13, e14))) | ~ (op2(e24, e20) = h5(op1(e14, e10))) | ~ (e23 = h5(op1(e14, e11))) | ~ (h5(op1(e14, e12)) = op2(e24, h5(e12))) | ~ (op2(e24, e23) = h5(op1(e14, e13))) | ~ (h5(e11) = h5(op1(e14, e14))) | sP66 | sP65 | sP64 | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f2864, f2472])).
fof(f2864, plain, (~ (h5(e12) = h5(op1(e13, e14))) | ~ (op2(e24, e20) = h5(op1(e14, e10))) | ~ (e23 = h5(op1(e14, e11))) | ~ (h5(op1(e14, e12)) = op2(e24, h5(e12))) | ~ (op2(e24, e23) = h5(op1(e14, e13))) | ~ (h5(e11) = h5(op1(e14, e14))) | sP66 | sP65 | sP64 | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f2863, f2473])).
fof(f2863, plain, (~ (op2(e23, e24) = h5(op1(e13, e14))) | ~ (op2(e24, e20) = h5(op1(e14, e10))) | ~ (e23 = h5(op1(e14, e11))) | ~ (h5(op1(e14, e12)) = op2(e24, h5(e12))) | ~ (op2(e24, e23) = h5(op1(e14, e13))) | ~ (h5(e11) = h5(op1(e14, e14))) | sP66 | sP65 | sP64 | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f2862, f2472])).
fof(f2862, plain, (~ (h5(op1(e13, e14)) = op2(h5(e13), e24)) | ~ (op2(e24, e20) = h5(op1(e14, e10))) | ~ (e23 = h5(op1(e14, e11))) | ~ (h5(op1(e14, e12)) = op2(e24, h5(e12))) | ~ (op2(e24, e23) = h5(op1(e14, e13))) | ~ (h5(e11) = h5(op1(e14, e14))) | sP66 | sP65 | sP64 | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f2861, f758])).
fof(f2861, plain, (~ (op2(e24, e20) = h5(op1(e14, e10))) | ~ (e23 = h5(op1(e14, e11))) | ~ (h5(op1(e14, e12)) = op2(e24, h5(e12))) | ~ (op2(e24, e23) = h5(op1(e14, e13))) | ~ (h5(e11) = h5(op1(e14, e14))) | sP66 | sP65 | sP64 | ~ (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f2860, f758])).
fof(f2860, plain, (~ (h5(op1(e14, e10)) = op2(h5(e14), e20)) | ~ (e23 = h5(op1(e14, e11))) | ~ (h5(op1(e14, e12)) = op2(e24, h5(e12))) | ~ (op2(e24, e23) = h5(op1(e14, e13))) | ~ (h5(e11) = h5(op1(e14, e14))) | sP66 | sP65 | sP64 | ~ (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f2859, f2492])).
fof(f2859, plain, (~ (e23 = h5(op1(e14, e11))) | ~ (h5(op1(e14, e12)) = op2(e24, h5(e12))) | ~ (op2(e24, e23) = h5(op1(e14, e13))) | ~ (h5(e11) = h5(op1(e14, e14))) | sP66 | sP65 | sP64 | ~ (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) | ~ (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f2858, f2489])).
fof(f2858, plain, (~ (h5(op1(e14, e11)) = op2(e24, h5(e11))) | ~ (h5(op1(e14, e12)) = op2(e24, h5(e12))) | ~ (op2(e24, e23) = h5(op1(e14, e13))) | ~ (h5(e11) = h5(op1(e14, e14))) | sP66 | sP65 | sP64 | ~ (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) | ~ (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f2857, f758])).
fof(f2857, plain, (~ (h5(op1(e14, e12)) = op2(e24, h5(e12))) | ~ (op2(e24, e23) = h5(op1(e14, e13))) | ~ (h5(e11) = h5(op1(e14, e14))) | sP66 | sP65 | sP64 | ~ (h5(op1(e14, e11)) = op2(h5(e14), h5(e11))) | ~ (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) | ~ (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f2856, f758])).
fof(f2856, plain, (~ (op2(e24, e23) = h5(op1(e14, e13))) | ~ (h5(e11) = h5(op1(e14, e14))) | sP66 | sP65 | sP64 | ~ (h5(op1(e14, e12)) = op2(h5(e14), h5(e12))) | ~ (h5(op1(e14, e11)) = op2(h5(e14), h5(e11))) | ~ (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) | ~ (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f2855, f758])).
fof(f2855, plain, (~ (h5(op1(e14, e13)) = op2(h5(e14), e23)) | ~ (h5(e11) = h5(op1(e14, e14))) | sP66 | sP65 | sP64 | ~ (h5(op1(e14, e12)) = op2(h5(e14), h5(e12))) | ~ (h5(op1(e14, e11)) = op2(h5(e14), h5(e11))) | ~ (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) | ~ (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f2854, f2472])).
fof(f2854, plain, (~ (h5(e11) = h5(op1(e14, e14))) | sP66 | sP65 | sP64 | ~ (h5(op1(e14, e13)) = op2(h5(e14), h5(e13))) | ~ (h5(op1(e14, e12)) = op2(h5(e14), h5(e12))) | ~ (h5(op1(e14, e11)) = op2(h5(e14), h5(e11))) | ~ (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) | ~ (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f2853, f760])).
fof(f2853, plain, (~ (op2(e24, e24) = h5(op1(e14, e14))) | sP66 | sP65 | sP64 | ~ (h5(op1(e14, e13)) = op2(h5(e14), h5(e13))) | ~ (h5(op1(e14, e12)) = op2(h5(e14), h5(e12))) | ~ (h5(op1(e14, e11)) = op2(h5(e14), h5(e11))) | ~ (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) | ~ (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f2852, f758])).
fof(f2852, plain, (sP66 | sP65 | sP64 | ~ (h5(op1(e14, e14)) = op2(h5(e14), h5(e14))) | ~ (h5(op1(e14, e13)) = op2(h5(e14), h5(e13))) | ~ (h5(op1(e14, e12)) = op2(h5(e14), h5(e12))) | ~ (h5(op1(e14, e11)) = op2(h5(e14), h5(e11))) | ~ (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) | ~ (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(subsumption_resolution, [], [f2851, f2494])).
fof(f2494, plain, ~ sP67, inference(subsumption_resolution, [], [f766, f2472])).
fof(f766, plain, (~ (e23 = h5(e13)) | ~ sP67), inference(cnf_transformation, [], [f143])).
fof(f143, plain, ((~ (e23 = h5(e14)) & ~ (e23 = h5(e13)) & ~ (e23 = h5(e12)) & ~ (e23 = h5(e11)) & ~ (e23 = h5(e10))) | ~ sP67), inference(nnf_transformation, [], [f93])).
fof(f2851, plain, (sP67 | sP66 | sP65 | sP64 | ~ (h5(op1(e14, e14)) = op2(h5(e14), h5(e14))) | ~ (h5(op1(e14, e13)) = op2(h5(e14), h5(e13))) | ~ (h5(op1(e14, e12)) = op2(h5(e14), h5(e12))) | ~ (h5(op1(e14, e11)) = op2(h5(e14), h5(e11))) | ~ (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) | ~ (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(subsumption_resolution, [], [f887, f758])).
fof(f887, plain, (~ (e24 = h5(e14)) | sP67 | sP66 | sP65 | sP64 | ~ (h5(op1(e14, e14)) = op2(h5(e14), h5(e14))) | ~ (h5(op1(e14, e13)) = op2(h5(e14), h5(e13))) | ~ (h5(op1(e14, e12)) = op2(h5(e14), h5(e12))) | ~ (h5(op1(e14, e11)) = op2(h5(e14), h5(e11))) | ~ (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) | ~ (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(cnf_transformation, [], [f94])).
fof(f2839, plain, (~ spl68_373 | ~ spl68_375), inference(avatar_split_clause, [], [f855, f2836, f2827])).
fof(f855, plain, (~ (e21 = h1(e12)) | ~ sP49), inference(cnf_transformation, [], [f161])).
fof(f161, plain, ((~ (e21 = h1(e14)) & ~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10))) | ~ sP49), inference(nnf_transformation, [], [f75])).
fof(f2824, plain, (~ spl68_368 | ~ spl68_372), inference(avatar_split_clause, [], [f848, f2821, f2802])).
fof(f848, plain, (~ (e22 = h1(e10)) | ~ sP50), inference(cnf_transformation, [], [f160])).
fof(f160, plain, ((~ (e22 = h1(e14)) & ~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10))) | ~ sP50), inference(nnf_transformation, [], [f76])).
fof(f2784, plain, (~ spl68_363 | ~ spl68_364), inference(avatar_split_clause, [], [f846, f2781, f2777])).
fof(f846, plain, (~ (e23 = h1(e13)) | ~ sP51), inference(cnf_transformation, [], [f159])).
fof(f159, plain, ((~ (e23 = h1(e14)) & ~ (e23 = h1(e13)) & ~ (e23 = h1(e12)) & ~ (e23 = h1(e11)) & ~ (e23 = h1(e10))) | ~ sP51), inference(nnf_transformation, [], [f77])).
fof(f2764, plain, (~ spl68_358 | ~ spl68_360), inference(avatar_split_clause, [], [f840, f2761, f2752])).
fof(f840, plain, (~ (e20 = h2(e12)) | ~ sP52), inference(cnf_transformation, [], [f158])).
fof(f158, plain, ((~ (e20 = h2(e14)) & ~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10))) | ~ sP52), inference(nnf_transformation, [], [f78])).
fof(f2759, plain, (~ spl68_358 | ~ spl68_359), inference(avatar_split_clause, [], [f841, f2756, f2752])).
fof(f841, plain, (~ (e20 = h2(e13)) | ~ sP52), inference(cnf_transformation, [], [f158])).
fof(f2743, plain, (~ spl68_353 | ~ spl68_356), inference(avatar_split_clause, [], [f829, f2740, f2726])).
fof(f829, plain, (~ (e22 = h2(e11)) | ~ sP54), inference(cnf_transformation, [], [f156])).
fof(f156, plain, ((~ (e22 = h2(e14)) & ~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10))) | ~ sP54), inference(nnf_transformation, [], [f80])).
fof(f2738, plain, (~ spl68_353 | ~ spl68_355), inference(avatar_split_clause, [], [f830, f2735, f2726])).
fof(f830, plain, (~ (e22 = h2(e12)) | ~ sP54), inference(cnf_transformation, [], [f156])).
fof(f2723, plain, (~ spl68_348 | ~ spl68_352), inference(avatar_split_clause, [], [f823, f2720, f2701])).
fof(f823, plain, (~ (e23 = h2(e10)) | ~ sP55), inference(cnf_transformation, [], [f155])).
fof(f155, plain, ((~ (e23 = h2(e14)) & ~ (e23 = h2(e13)) & ~ (e23 = h2(e12)) & ~ (e23 = h2(e11)) & ~ (e23 = h2(e10))) | ~ sP55), inference(nnf_transformation, [], [f81])).
fof(f2708, plain, (~ spl68_348 | ~ spl68_349), inference(avatar_split_clause, [], [f826, f2705, f2701])).
fof(f826, plain, (~ (e23 = h2(e13)) | ~ sP55), inference(cnf_transformation, [], [f155])).
fof(f2688, plain, (~ spl68_343 | ~ spl68_345), inference(avatar_split_clause, [], [f820, f2685, f2676])).
fof(f820, plain, (~ (e20 = h3(e12)) | ~ sP56), inference(cnf_transformation, [], [f154])).
fof(f154, plain, ((~ (e20 = h3(e14)) & ~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10))) | ~ sP56), inference(nnf_transformation, [], [f82])).
fof(f2668, plain, (~ spl68_338 | ~ spl68_341), inference(avatar_split_clause, [], [f814, f2665, f2651])).
fof(f814, plain, (~ (e21 = h3(e11)) | ~ sP57), inference(cnf_transformation, [], [f153])).
fof(f153, plain, ((~ (e21 = h3(e14)) & ~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10))) | ~ sP57), inference(nnf_transformation, [], [f83])).
fof(f2647, plain, (~ spl68_333 | ~ spl68_337), inference(avatar_split_clause, [], [f803, f2644, f2625])).
fof(f803, plain, (~ (e23 = h3(e10)) | ~ sP59), inference(cnf_transformation, [], [f151])).
fof(f151, plain, ((~ (e23 = h3(e14)) & ~ (e23 = h3(e13)) & ~ (e23 = h3(e12)) & ~ (e23 = h3(e11)) & ~ (e23 = h3(e10))) | ~ sP59), inference(nnf_transformation, [], [f85])).
fof(f2607, plain, (~ spl68_328 | ~ spl68_329), inference(avatar_split_clause, [], [f801, f2604, f2600])).
fof(f801, plain, (~ (e20 = h4(e13)) | ~ sP60), inference(cnf_transformation, [], [f150])).
fof(f150, plain, ((~ (e20 = h4(e14)) & ~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10))) | ~ sP60), inference(nnf_transformation, [], [f86])).
fof(f2592, plain, (~ spl68_323 | ~ spl68_326), inference(avatar_split_clause, [], [f794, f2589, f2575])).
fof(f794, plain, (~ (e21 = h4(e11)) | ~ sP61), inference(cnf_transformation, [], [f149])).
fof(f149, plain, ((~ (e21 = h4(e14)) & ~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10))) | ~ sP61), inference(nnf_transformation, [], [f87])).
fof(f2572, plain, (~ spl68_318 | ~ spl68_322), inference(avatar_split_clause, [], [f788, f2569, f2550])).
fof(f788, plain, (~ (e22 = h4(e10)) | ~ sP62), inference(cnf_transformation, [], [f148])).
fof(f148, plain, ((~ (e22 = h4(e14)) & ~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10))) | ~ sP62), inference(nnf_transformation, [], [f88])).
fof(f2546, plain, ~ spl68_315, inference(avatar_split_clause, [], [f2545, f2532])).
fof(f2545, plain, ~ sP64, inference(subsumption_resolution, [], [f778, f2492])).
fof(f778, plain, (~ (e20 = h5(e10)) | ~ sP64), inference(cnf_transformation, [], [f146])).
fof(f146, plain, ((~ (e20 = h5(e14)) & ~ (e20 = h5(e13)) & ~ (e20 = h5(e12)) & ~ (e20 = h5(e11)) & ~ (e20 = h5(e10))) | ~ sP64), inference(nnf_transformation, [], [f90])).
fof(f2527, plain, (~ spl68_312 | ~ spl68_314), inference(avatar_split_clause, [], [f774, f2524, f2515])).
fof(f774, plain, (~ (e21 = h5(e11)) | ~ sP65), inference(cnf_transformation, [], [f145])).
fof(f145, plain, ((~ (e21 = h5(e14)) & ~ (e21 = h5(e13)) & ~ (e21 = h5(e12)) & ~ (e21 = h5(e11)) & ~ (e21 = h5(e10))) | ~ sP65), inference(nnf_transformation, [], [f91])).
fof(f2505, plain, (~ spl68_309 | ~ spl68_310), inference(avatar_split_clause, [], [f770, f2502, f2498])).
fof(f770, plain, (~ (e22 = h5(e12)) | ~ sP66), inference(cnf_transformation, [], [f144])).
fof(f144, plain, ((~ (e22 = h5(e14)) & ~ (e22 = h5(e13)) & ~ (e22 = h5(e12)) & ~ (e22 = h5(e11)) & ~ (e22 = h5(e10))) | ~ sP66), inference(nnf_transformation, [], [f92])).
fof(f2427, plain, spl68_132, inference(avatar_split_clause, [], [f735, f1489])).
fof(f735, plain, (e21 = op2(e24, e24)), inference(cnf_transformation, [], [f15])).
fof(f2426, plain, spl68_158, inference(avatar_split_clause, [], [f2425, f1598])).
fof(f2425, plain, (e22 = op2(e23, e24)), inference(forward_demodulation, [], [f736, f737])).
fof(f736, plain, (e22 = op2(op2(e24, op2(e24, e24)), e24)), inference(cnf_transformation, [], [f15])).
fof(f2424, plain, spl68_2, inference(avatar_split_clause, [], [f731, f893])).
fof(f731, plain, (e11 = op1(e14, e14)), inference(cnf_transformation, [], [f14])).
fof(f2423, plain, spl68_28, inference(avatar_split_clause, [], [f2422, f1002])).
fof(f2422, plain, (e12 = op1(e13, e14)), inference(forward_demodulation, [], [f732, f733])).
fof(f732, plain, (e12 = op1(op1(e14, op1(e14, e14)), e14)), inference(cnf_transformation, [], [f14])).
fof(f2421, plain, (spl68_308 | spl68_307 | spl68_306 | spl68_305 | spl68_304 | spl68_303 | spl68_302 | spl68_301 | spl68_300 | spl68_299 | spl68_298 | spl68_297 | spl68_296 | spl68_295 | spl68_294 | spl68_293 | spl68_292 | spl68_291 | spl68_290 | spl68_289 | spl68_288 | spl68_287 | spl68_286 | spl68_285 | spl68_135), inference(avatar_split_clause, [], [f727, f1501, f2252, f2259, f2266, f2273, f2280, f2287, f2294, f2301, f2308, f2315, f2322, f2329, f2336, f2343, f2350, f2357, f2364, f2371, f2378, f2385, f2392, f2399, f2406, f2413])).
fof(f2413, plain, (spl68_308 <=> sP24), introduced(avatar_definition, [new_symbols(naming, [spl68_308])])).
fof(f2406, plain, (spl68_307 <=> sP25), introduced(avatar_definition, [new_symbols(naming, [spl68_307])])).
fof(f2399, plain, (spl68_306 <=> sP26), introduced(avatar_definition, [new_symbols(naming, [spl68_306])])).
fof(f2392, plain, (spl68_305 <=> sP27), introduced(avatar_definition, [new_symbols(naming, [spl68_305])])).
fof(f2385, plain, (spl68_304 <=> sP28), introduced(avatar_definition, [new_symbols(naming, [spl68_304])])).
fof(f2378, plain, (spl68_303 <=> sP29), introduced(avatar_definition, [new_symbols(naming, [spl68_303])])).
fof(f2371, plain, (spl68_302 <=> sP30), introduced(avatar_definition, [new_symbols(naming, [spl68_302])])).
fof(f2364, plain, (spl68_301 <=> sP31), introduced(avatar_definition, [new_symbols(naming, [spl68_301])])).
fof(f2357, plain, (spl68_300 <=> sP32), introduced(avatar_definition, [new_symbols(naming, [spl68_300])])).
fof(f2350, plain, (spl68_299 <=> sP33), introduced(avatar_definition, [new_symbols(naming, [spl68_299])])).
fof(f2343, plain, (spl68_298 <=> sP34), introduced(avatar_definition, [new_symbols(naming, [spl68_298])])).
fof(f2336, plain, (spl68_297 <=> sP35), introduced(avatar_definition, [new_symbols(naming, [spl68_297])])).
fof(f2329, plain, (spl68_296 <=> sP36), introduced(avatar_definition, [new_symbols(naming, [spl68_296])])).
fof(f2322, plain, (spl68_295 <=> sP37), introduced(avatar_definition, [new_symbols(naming, [spl68_295])])).
fof(f2315, plain, (spl68_294 <=> sP38), introduced(avatar_definition, [new_symbols(naming, [spl68_294])])).
fof(f2308, plain, (spl68_293 <=> sP39), introduced(avatar_definition, [new_symbols(naming, [spl68_293])])).
fof(f2301, plain, (spl68_292 <=> sP40), introduced(avatar_definition, [new_symbols(naming, [spl68_292])])).
fof(f2294, plain, (spl68_291 <=> sP41), introduced(avatar_definition, [new_symbols(naming, [spl68_291])])).
fof(f2287, plain, (spl68_290 <=> sP42), introduced(avatar_definition, [new_symbols(naming, [spl68_290])])).
fof(f2280, plain, (spl68_289 <=> sP43), introduced(avatar_definition, [new_symbols(naming, [spl68_289])])).
fof(f2273, plain, (spl68_288 <=> sP44), introduced(avatar_definition, [new_symbols(naming, [spl68_288])])).
fof(f2266, plain, (spl68_287 <=> sP45), introduced(avatar_definition, [new_symbols(naming, [spl68_287])])).
fof(f2259, plain, (spl68_286 <=> sP46), introduced(avatar_definition, [new_symbols(naming, [spl68_286])])).
fof(f2252, plain, (spl68_285 <=> sP47), introduced(avatar_definition, [new_symbols(naming, [spl68_285])])).
fof(f1501, plain, (spl68_135 <=> (e24 = op2(e24, e24))), introduced(avatar_definition, [new_symbols(naming, [spl68_135])])).
fof(f727, plain, ((e24 = op2(e24, e24)) | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24), inference(cnf_transformation, [], [f73])).
fof(f73, plain, ((~ (e24 = op2(e24, e24)) & (e24 = op2(e24, e24)) & (e24 = op2(e24, e24))) | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24), inference(definition_folding, [], [f13, e72, e71, e70, e69, e68, e67, e66, e65, e64, e63, e62, e61, e60, e59, e58, e57, e56, e55, e54, e53, e52, e51, e50, e49])).
fof(f49, plain, ((~ (e20 = op2(e20, e20)) & (e20 = op2(e20, e20)) & (e20 = op2(e20, e20))) | ~ sP24), inference(usedef, [], [e49])).
fof(e49, plain, (sP24 <=> (~ (e20 = op2(e20, e20)) & (e20 = op2(e20, e20)) & (e20 = op2(e20, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP24])])).
fof(f50, plain, ((~ (e20 = op2(e20, e21)) & (e20 = op2(e21, e21)) & (op2(e20, e20) = e21)) | ~ sP25), inference(usedef, [], [e50])).
fof(e50, plain, (sP25 <=> (~ (e20 = op2(e20, e21)) & (e20 = op2(e21, e21)) & (op2(e20, e20) = e21))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP25])])).
fof(f51, plain, ((~ (e20 = op2(e20, e22)) & (e20 = op2(e22, e22)) & (op2(e20, e20) = e22)) | ~ sP26), inference(usedef, [], [e51])).
fof(e51, plain, (sP26 <=> (~ (e20 = op2(e20, e22)) & (e20 = op2(e22, e22)) & (op2(e20, e20) = e22))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP26])])).
fof(f52, plain, ((~ (e20 = op2(e20, e23)) & (e20 = op2(e23, e23)) & (op2(e20, e20) = e23)) | ~ sP27), inference(usedef, [], [e52])).
fof(e52, plain, (sP27 <=> (~ (e20 = op2(e20, e23)) & (e20 = op2(e23, e23)) & (op2(e20, e20) = e23))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP27])])).
fof(f53, plain, ((~ (e20 = op2(e20, e24)) & (e20 = op2(e24, e24)) & (op2(e20, e20) = e24)) | ~ sP28), inference(usedef, [], [e53])).
fof(e53, plain, (sP28 <=> (~ (e20 = op2(e20, e24)) & (e20 = op2(e24, e24)) & (op2(e20, e20) = e24))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP28])])).
fof(f54, plain, ((~ (e21 = op2(e21, e20)) & (op2(e20, e20) = e21) & (e20 = op2(e21, e21))) | ~ sP29), inference(usedef, [], [e54])).
fof(e54, plain, (sP29 <=> (~ (e21 = op2(e21, e20)) & (op2(e20, e20) = e21) & (e20 = op2(e21, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP29])])).
fof(f55, plain, ((~ (e21 = op2(e21, e21)) & (e21 = op2(e21, e21)) & (e21 = op2(e21, e21))) | ~ sP30), inference(usedef, [], [e55])).
fof(e55, plain, (sP30 <=> (~ (e21 = op2(e21, e21)) & (e21 = op2(e21, e21)) & (e21 = op2(e21, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP30])])).
fof(f56, plain, ((~ (e21 = op2(e21, e22)) & (e21 = op2(e22, e22)) & (e22 = op2(e21, e21))) | ~ sP31), inference(usedef, [], [e56])).
fof(e56, plain, (sP31 <=> (~ (e21 = op2(e21, e22)) & (e21 = op2(e22, e22)) & (e22 = op2(e21, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP31])])).
fof(f57, plain, ((~ (e21 = op2(e21, e23)) & (e21 = op2(e23, e23)) & (e23 = op2(e21, e21))) | ~ sP32), inference(usedef, [], [e57])).
fof(e57, plain, (sP32 <=> (~ (e21 = op2(e21, e23)) & (e21 = op2(e23, e23)) & (e23 = op2(e21, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP32])])).
fof(f58, plain, ((~ (e21 = op2(e21, e24)) & (e21 = op2(e24, e24)) & (e24 = op2(e21, e21))) | ~ sP33), inference(usedef, [], [e58])).
fof(e58, plain, (sP33 <=> (~ (e21 = op2(e21, e24)) & (e21 = op2(e24, e24)) & (e24 = op2(e21, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP33])])).
fof(f59, plain, ((~ (e22 = op2(e22, e20)) & (op2(e20, e20) = e22) & (e20 = op2(e22, e22))) | ~ sP34), inference(usedef, [], [e59])).
fof(e59, plain, (sP34 <=> (~ (e22 = op2(e22, e20)) & (op2(e20, e20) = e22) & (e20 = op2(e22, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP34])])).
fof(f60, plain, ((~ (e22 = op2(e22, e21)) & (e22 = op2(e21, e21)) & (e21 = op2(e22, e22))) | ~ sP35), inference(usedef, [], [e60])).
fof(e60, plain, (sP35 <=> (~ (e22 = op2(e22, e21)) & (e22 = op2(e21, e21)) & (e21 = op2(e22, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP35])])).
fof(f61, plain, ((~ (e22 = op2(e22, e22)) & (e22 = op2(e22, e22)) & (e22 = op2(e22, e22))) | ~ sP36), inference(usedef, [], [e61])).
fof(e61, plain, (sP36 <=> (~ (e22 = op2(e22, e22)) & (e22 = op2(e22, e22)) & (e22 = op2(e22, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP36])])).
fof(f62, plain, ((~ (e22 = op2(e22, e23)) & (e22 = op2(e23, e23)) & (e23 = op2(e22, e22))) | ~ sP37), inference(usedef, [], [e62])).
fof(e62, plain, (sP37 <=> (~ (e22 = op2(e22, e23)) & (e22 = op2(e23, e23)) & (e23 = op2(e22, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP37])])).
fof(f63, plain, ((~ (e22 = op2(e22, e24)) & (e22 = op2(e24, e24)) & (e24 = op2(e22, e22))) | ~ sP38), inference(usedef, [], [e63])).
fof(e63, plain, (sP38 <=> (~ (e22 = op2(e22, e24)) & (e22 = op2(e24, e24)) & (e24 = op2(e22, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP38])])).
fof(f64, plain, ((~ (e23 = op2(e23, e20)) & (op2(e20, e20) = e23) & (e20 = op2(e23, e23))) | ~ sP39), inference(usedef, [], [e64])).
fof(e64, plain, (sP39 <=> (~ (e23 = op2(e23, e20)) & (op2(e20, e20) = e23) & (e20 = op2(e23, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP39])])).
fof(f65, plain, ((~ (e23 = op2(e23, e21)) & (e23 = op2(e21, e21)) & (e21 = op2(e23, e23))) | ~ sP40), inference(usedef, [], [e65])).
fof(e65, plain, (sP40 <=> (~ (e23 = op2(e23, e21)) & (e23 = op2(e21, e21)) & (e21 = op2(e23, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP40])])).
fof(f66, plain, ((~ (e23 = op2(e23, e22)) & (e23 = op2(e22, e22)) & (e22 = op2(e23, e23))) | ~ sP41), inference(usedef, [], [e66])).
fof(e66, plain, (sP41 <=> (~ (e23 = op2(e23, e22)) & (e23 = op2(e22, e22)) & (e22 = op2(e23, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP41])])).
fof(f67, plain, ((~ (e23 = op2(e23, e23)) & (e23 = op2(e23, e23)) & (e23 = op2(e23, e23))) | ~ sP42), inference(usedef, [], [e67])).
fof(e67, plain, (sP42 <=> (~ (e23 = op2(e23, e23)) & (e23 = op2(e23, e23)) & (e23 = op2(e23, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP42])])).
fof(f68, plain, ((~ (e23 = op2(e23, e24)) & (e23 = op2(e24, e24)) & (e24 = op2(e23, e23))) | ~ sP43), inference(usedef, [], [e68])).
fof(e68, plain, (sP43 <=> (~ (e23 = op2(e23, e24)) & (e23 = op2(e24, e24)) & (e24 = op2(e23, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP43])])).
fof(f69, plain, ((~ (e24 = op2(e24, e20)) & (op2(e20, e20) = e24) & (e20 = op2(e24, e24))) | ~ sP44), inference(usedef, [], [e69])).
fof(e69, plain, (sP44 <=> (~ (e24 = op2(e24, e20)) & (op2(e20, e20) = e24) & (e20 = op2(e24, e24)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP44])])).
fof(f70, plain, ((~ (e24 = op2(e24, e21)) & (e24 = op2(e21, e21)) & (e21 = op2(e24, e24))) | ~ sP45), inference(usedef, [], [e70])).
fof(e70, plain, (sP45 <=> (~ (e24 = op2(e24, e21)) & (e24 = op2(e21, e21)) & (e21 = op2(e24, e24)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP45])])).
fof(f71, plain, ((~ (e24 = op2(e24, e22)) & (e24 = op2(e22, e22)) & (e22 = op2(e24, e24))) | ~ sP46), inference(usedef, [], [e71])).
fof(e71, plain, (sP46 <=> (~ (e24 = op2(e24, e22)) & (e24 = op2(e22, e22)) & (e22 = op2(e24, e24)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP46])])).
fof(f72, plain, ((~ (e24 = op2(e24, e23)) & (e24 = op2(e23, e23)) & (e23 = op2(e24, e24))) | ~ sP47), inference(usedef, [], [e72])).
fof(e72, plain, (sP47 <=> (~ (e24 = op2(e24, e23)) & (e24 = op2(e23, e23)) & (e23 = op2(e24, e24)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP47])])).
fof(f13, plain, ((~ (e24 = op2(e24, e24)) & (e24 = op2(e24, e24)) & (e24 = op2(e24, e24))) | (~ (e24 = op2(e24, e23)) & (e24 = op2(e23, e23)) & (e23 = op2(e24, e24))) | (~ (e24 = op2(e24, e22)) & (e24 = op2(e22, e22)) & (e22 = op2(e24, e24))) | (~ (e24 = op2(e24, e21)) & (e24 = op2(e21, e21)) & (e21 = op2(e24, e24))) | (~ (e24 = op2(e24, e20)) & (op2(e20, e20) = e24) & (e20 = op2(e24, e24))) | (~ (e23 = op2(e23, e24)) & (e23 = op2(e24, e24)) & (e24 = op2(e23, e23))) | (~ (e23 = op2(e23, e23)) & (e23 = op2(e23, e23)) & (e23 = op2(e23, e23))) | (~ (e23 = op2(e23, e22)) & (e23 = op2(e22, e22)) & (e22 = op2(e23, e23))) | (~ (e23 = op2(e23, e21)) & (e23 = op2(e21, e21)) & (e21 = op2(e23, e23))) | (~ (e23 = op2(e23, e20)) & (op2(e20, e20) = e23) & (e20 = op2(e23, e23))) | (~ (e22 = op2(e22, e24)) & (e22 = op2(e24, e24)) & (e24 = op2(e22, e22))) | (~ (e22 = op2(e22, e23)) & (e22 = op2(e23, e23)) & (e23 = op2(e22, e22))) | (~ (e22 = op2(e22, e22)) & (e22 = op2(e22, e22)) & (e22 = op2(e22, e22))) | (~ (e22 = op2(e22, e21)) & (e22 = op2(e21, e21)) & (e21 = op2(e22, e22))) | (~ (e22 = op2(e22, e20)) & (op2(e20, e20) = e22) & (e20 = op2(e22, e22))) | (~ (e21 = op2(e21, e24)) & (e21 = op2(e24, e24)) & (e24 = op2(e21, e21))) | (~ (e21 = op2(e21, e23)) & (e21 = op2(e23, e23)) & (e23 = op2(e21, e21))) | (~ (e21 = op2(e21, e22)) & (e21 = op2(e22, e22)) & (e22 = op2(e21, e21))) | (~ (e21 = op2(e21, e21)) & (e21 = op2(e21, e21)) & (e21 = op2(e21, e21))) | (~ (e21 = op2(e21, e20)) & (op2(e20, e20) = e21) & (e20 = op2(e21, e21))) | (~ (e20 = op2(e20, e24)) & (e20 = op2(e24, e24)) & (op2(e20, e20) = e24)) | (~ (e20 = op2(e20, e23)) & (e20 = op2(e23, e23)) & (op2(e20, e20) = e23)) | (~ (e20 = op2(e20, e22)) & (e20 = op2(e22, e22)) & (op2(e20, e20) = e22)) | (~ (e20 = op2(e20, e21)) & (e20 = op2(e21, e21)) & (op2(e20, e20) = e21)) | (~ (e20 = op2(e20, e20)) & (e20 = op2(e20, e20)) & (e20 = op2(e20, e20)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG052+1.p', ax13)).
fof(f2419, plain, (spl68_308 | spl68_307 | spl68_306 | spl68_305 | spl68_304 | spl68_303 | spl68_302 | spl68_301 | spl68_300 | spl68_299 | spl68_298 | spl68_297 | spl68_296 | spl68_295 | spl68_294 | spl68_293 | spl68_292 | spl68_291 | spl68_290 | spl68_289 | spl68_288 | spl68_287 | spl68_286 | spl68_285 | ~ spl68_135), inference(avatar_split_clause, [], [f729, f1501, f2252, f2259, f2266, f2273, f2280, f2287, f2294, f2301, f2308, f2315, f2322, f2329, f2336, f2343, f2350, f2357, f2364, f2371, f2378, f2385, f2392, f2399, f2406, f2413])).
fof(f729, plain, (~ (e24 = op2(e24, e24)) | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24), inference(cnf_transformation, [], [f73])).
fof(f2418, plain, (~ spl68_308 | spl68_251), inference(avatar_split_clause, [], [f724, f1989, f2413])).
fof(f724, plain, ((e20 = op2(e20, e20)) | ~ sP24), inference(cnf_transformation, [], [f142])).
fof(f142, plain, ((~ (e20 = op2(e20, e20)) & (e20 = op2(e20, e20)) & (e20 = op2(e20, e20))) | ~ sP24), inference(nnf_transformation, [], [f49])).
fof(f2416, plain, (~ spl68_308 | ~ spl68_251), inference(avatar_split_clause, [], [f726, f1989, f2413])).
fof(f726, plain, (~ (e20 = op2(e20, e20)) | ~ sP24), inference(cnf_transformation, [], [f142])).
fof(f2410, plain, (~ spl68_307 | spl68_221), inference(avatar_split_clause, [], [f722, f1863, f2406])).
fof(f722, plain, ((e20 = op2(e21, e21)) | ~ sP25), inference(cnf_transformation, [], [f141])).
fof(f141, plain, ((~ (e20 = op2(e20, e21)) & (e20 = op2(e21, e21)) & (op2(e20, e20) = e21)) | ~ sP25), inference(nnf_transformation, [], [f50])).
fof(f2404, plain, (~ spl68_306 | spl68_253), inference(avatar_split_clause, [], [f718, f1997, f2399])).
fof(f718, plain, ((op2(e20, e20) = e22) | ~ sP26), inference(cnf_transformation, [], [f140])).
fof(f140, plain, ((~ (e20 = op2(e20, e22)) & (e20 = op2(e22, e22)) & (op2(e20, e20) = e22)) | ~ sP26), inference(nnf_transformation, [], [f51])).
fof(f2402, plain, (~ spl68_306 | ~ spl68_241), inference(avatar_split_clause, [], [f720, f1947, f2399])).
fof(f720, plain, (~ (e20 = op2(e20, e22)) | ~ sP26), inference(cnf_transformation, [], [f140])).
fof(f2397, plain, (~ spl68_305 | spl68_254), inference(avatar_split_clause, [], [f715, f2001, f2392])).
fof(f715, plain, ((op2(e20, e20) = e23) | ~ sP27), inference(cnf_transformation, [], [f139])).
fof(f139, plain, ((~ (e20 = op2(e20, e23)) & (e20 = op2(e23, e23)) & (op2(e20, e20) = e23)) | ~ sP27), inference(nnf_transformation, [], [f52])).
fof(f2389, plain, (~ spl68_304 | spl68_131), inference(avatar_split_clause, [], [f713, f1485, f2385])).
fof(f713, plain, ((e20 = op2(e24, e24)) | ~ sP28), inference(cnf_transformation, [], [f138])).
fof(f138, plain, ((~ (e20 = op2(e20, e24)) & (e20 = op2(e24, e24)) & (op2(e20, e20) = e24)) | ~ sP28), inference(nnf_transformation, [], [f53])).
fof(f2383, plain, (~ spl68_303 | spl68_221), inference(avatar_split_clause, [], [f709, f1863, f2378])).
fof(f709, plain, ((e20 = op2(e21, e21)) | ~ sP29), inference(cnf_transformation, [], [f137])).
fof(f137, plain, ((~ (e21 = op2(e21, e20)) & (op2(e20, e20) = e21) & (e20 = op2(e21, e21))) | ~ sP29), inference(nnf_transformation, [], [f54])).
fof(f2376, plain, (~ spl68_302 | spl68_222), inference(avatar_split_clause, [], [f706, f1867, f2371])).
fof(f706, plain, ((e21 = op2(e21, e21)) | ~ sP30), inference(cnf_transformation, [], [f136])).
fof(f136, plain, ((~ (e21 = op2(e21, e21)) & (e21 = op2(e21, e21)) & (e21 = op2(e21, e21))) | ~ sP30), inference(nnf_transformation, [], [f55])).
fof(f2368, plain, (~ spl68_301 | spl68_192), inference(avatar_split_clause, [], [f704, f1741, f2364])).
fof(f704, plain, ((e21 = op2(e22, e22)) | ~ sP31), inference(cnf_transformation, [], [f135])).
fof(f135, plain, ((~ (e21 = op2(e21, e22)) & (e21 = op2(e22, e22)) & (e22 = op2(e21, e21))) | ~ sP31), inference(nnf_transformation, [], [f56])).
fof(f2367, plain, (~ spl68_301 | ~ spl68_217), inference(avatar_split_clause, [], [f705, f1846, f2364])).
fof(f705, plain, (~ (e21 = op2(e21, e22)) | ~ sP31), inference(cnf_transformation, [], [f135])).
fof(f2362, plain, (~ spl68_300 | spl68_224), inference(avatar_split_clause, [], [f700, f1875, f2357])).
fof(f700, plain, ((e23 = op2(e21, e21)) | ~ sP32), inference(cnf_transformation, [], [f134])).
fof(f134, plain, ((~ (e21 = op2(e21, e23)) & (e21 = op2(e23, e23)) & (e23 = op2(e21, e21))) | ~ sP32), inference(nnf_transformation, [], [f57])).
fof(f2355, plain, (~ spl68_299 | spl68_225), inference(avatar_split_clause, [], [f697, f1879, f2350])).
fof(f697, plain, ((e24 = op2(e21, e21)) | ~ sP33), inference(cnf_transformation, [], [f133])).
fof(f133, plain, ((~ (e21 = op2(e21, e24)) & (e21 = op2(e24, e24)) & (e24 = op2(e21, e21))) | ~ sP33), inference(nnf_transformation, [], [f58])).
fof(f2348, plain, (~ spl68_298 | spl68_191), inference(avatar_split_clause, [], [f694, f1737, f2343])).
fof(f694, plain, ((e20 = op2(e22, e22)) | ~ sP34), inference(cnf_transformation, [], [f132])).
fof(f132, plain, ((~ (e22 = op2(e22, e20)) & (op2(e20, e20) = e22) & (e20 = op2(e22, e22))) | ~ sP34), inference(nnf_transformation, [], [f59])).
fof(f2347, plain, (~ spl68_298 | spl68_253), inference(avatar_split_clause, [], [f695, f1997, f2343])).
fof(f695, plain, ((op2(e20, e20) = e22) | ~ sP34), inference(cnf_transformation, [], [f132])).
fof(f2341, plain, (~ spl68_297 | spl68_192), inference(avatar_split_clause, [], [f691, f1741, f2336])).
fof(f691, plain, ((e21 = op2(e22, e22)) | ~ sP35), inference(cnf_transformation, [], [f131])).
fof(f131, plain, ((~ (e22 = op2(e22, e21)) & (e22 = op2(e21, e21)) & (e21 = op2(e22, e22))) | ~ sP35), inference(nnf_transformation, [], [f60])).
fof(f2334, plain, (~ spl68_296 | spl68_193), inference(avatar_split_clause, [], [f688, f1745, f2329])).
fof(f688, plain, ((e22 = op2(e22, e22)) | ~ sP36), inference(cnf_transformation, [], [f130])).
fof(f130, plain, ((~ (e22 = op2(e22, e22)) & (e22 = op2(e22, e22)) & (e22 = op2(e22, e22))) | ~ sP36), inference(nnf_transformation, [], [f61])).
fof(f2332, plain, (~ spl68_296 | ~ spl68_193), inference(avatar_split_clause, [], [f690, f1745, f2329])).
fof(f690, plain, (~ (e22 = op2(e22, e22)) | ~ sP36), inference(cnf_transformation, [], [f130])).
fof(f2326, plain, (~ spl68_295 | spl68_163), inference(avatar_split_clause, [], [f686, f1619, f2322])).
fof(f686, plain, ((e22 = op2(e23, e23)) | ~ sP37), inference(cnf_transformation, [], [f129])).
fof(f129, plain, ((~ (e22 = op2(e22, e23)) & (e22 = op2(e23, e23)) & (e23 = op2(e22, e22))) | ~ sP37), inference(nnf_transformation, [], [f62])).
fof(f2319, plain, (~ spl68_294 | spl68_133), inference(avatar_split_clause, [], [f683, f1493, f2315])).
fof(f683, plain, ((e22 = op2(e24, e24)) | ~ sP38), inference(cnf_transformation, [], [f128])).
fof(f128, plain, ((~ (e22 = op2(e22, e24)) & (e22 = op2(e24, e24)) & (e24 = op2(e22, e22))) | ~ sP38), inference(nnf_transformation, [], [f63])).
fof(f2312, plain, (~ spl68_293 | spl68_254), inference(avatar_split_clause, [], [f680, f2001, f2308])).
fof(f680, plain, ((op2(e20, e20) = e23) | ~ sP39), inference(cnf_transformation, [], [f127])).
fof(f127, plain, ((~ (e23 = op2(e23, e20)) & (op2(e20, e20) = e23) & (e20 = op2(e23, e23))) | ~ sP39), inference(nnf_transformation, [], [f64])).
fof(f2311, plain, (~ spl68_293 | ~ spl68_179), inference(avatar_split_clause, [], [f681, f1686, f2308])).
fof(f681, plain, (~ (e23 = op2(e23, e20)) | ~ sP39), inference(cnf_transformation, [], [f127])).
fof(f2305, plain, (~ spl68_292 | spl68_224), inference(avatar_split_clause, [], [f677, f1875, f2301])).
fof(f677, plain, ((e23 = op2(e21, e21)) | ~ sP40), inference(cnf_transformation, [], [f126])).
fof(f126, plain, ((~ (e23 = op2(e23, e21)) & (e23 = op2(e21, e21)) & (e21 = op2(e23, e23))) | ~ sP40), inference(nnf_transformation, [], [f65])).
fof(f2299, plain, (~ spl68_291 | spl68_163), inference(avatar_split_clause, [], [f673, f1619, f2294])).
fof(f673, plain, ((e22 = op2(e23, e23)) | ~ sP41), inference(cnf_transformation, [], [f125])).
fof(f125, plain, ((~ (e23 = op2(e23, e22)) & (e23 = op2(e22, e22)) & (e22 = op2(e23, e23))) | ~ sP41), inference(nnf_transformation, [], [f66])).
fof(f2292, plain, (~ spl68_290 | spl68_164), inference(avatar_split_clause, [], [f670, f1623, f2287])).
fof(f670, plain, ((e23 = op2(e23, e23)) | ~ sP42), inference(cnf_transformation, [], [f124])).
fof(f124, plain, ((~ (e23 = op2(e23, e23)) & (e23 = op2(e23, e23)) & (e23 = op2(e23, e23))) | ~ sP42), inference(nnf_transformation, [], [f67])).
fof(f2284, plain, (~ spl68_289 | spl68_134), inference(avatar_split_clause, [], [f668, f1497, f2280])).
fof(f668, plain, ((e23 = op2(e24, e24)) | ~ sP43), inference(cnf_transformation, [], [f123])).
fof(f123, plain, ((~ (e23 = op2(e23, e24)) & (e23 = op2(e24, e24)) & (e24 = op2(e23, e23))) | ~ sP43), inference(nnf_transformation, [], [f68])).
fof(f2278, plain, (~ spl68_288 | spl68_131), inference(avatar_split_clause, [], [f664, f1485, f2273])).
fof(f664, plain, ((e20 = op2(e24, e24)) | ~ sP44), inference(cnf_transformation, [], [f122])).
fof(f122, plain, ((~ (e24 = op2(e24, e20)) & (op2(e20, e20) = e24) & (e20 = op2(e24, e24))) | ~ sP44), inference(nnf_transformation, [], [f69])).
fof(f2270, plain, (~ spl68_287 | spl68_225), inference(avatar_split_clause, [], [f662, f1879, f2266])).
fof(f662, plain, ((e24 = op2(e21, e21)) | ~ sP45), inference(cnf_transformation, [], [f121])).
fof(f121, plain, ((~ (e24 = op2(e24, e21)) & (e24 = op2(e21, e21)) & (e21 = op2(e24, e24))) | ~ sP45), inference(nnf_transformation, [], [f70])).
fof(f2264, plain, (~ spl68_286 | spl68_133), inference(avatar_split_clause, [], [f658, f1493, f2259])).
fof(f658, plain, ((e22 = op2(e24, e24)) | ~ sP46), inference(cnf_transformation, [], [f120])).
fof(f120, plain, ((~ (e24 = op2(e24, e22)) & (e24 = op2(e22, e22)) & (e22 = op2(e24, e24))) | ~ sP46), inference(nnf_transformation, [], [f71])).
fof(f2257, plain, (~ spl68_285 | spl68_134), inference(avatar_split_clause, [], [f655, f1497, f2252])).
fof(f655, plain, ((e23 = op2(e24, e24)) | ~ sP47), inference(cnf_transformation, [], [f119])).
fof(f119, plain, ((~ (e24 = op2(e24, e23)) & (e24 = op2(e23, e23)) & (e23 = op2(e24, e24))) | ~ sP47), inference(nnf_transformation, [], [f72])).
fof(f2250, plain, (spl68_284 | spl68_283 | spl68_282 | spl68_281 | spl68_280 | spl68_279 | spl68_278 | spl68_277 | spl68_276 | spl68_275 | spl68_274 | spl68_273 | spl68_272 | spl68_271 | spl68_270 | spl68_269 | spl68_268 | spl68_267 | spl68_266 | spl68_265 | spl68_264 | spl68_263 | spl68_262 | spl68_261 | spl68_5), inference(avatar_split_clause, [], [f652, f905, f2081, f2088, f2095, f2102, f2109, f2116, f2123, f2130, f2137, f2144, f2151, f2158, f2165, f2172, f2179, f2186, f2193, f2200, f2207, f2214, f2221, f2228, f2235, f2242])).
fof(f2242, plain, (spl68_284 <=> sP0), introduced(avatar_definition, [new_symbols(naming, [spl68_284])])).
fof(f2235, plain, (spl68_283 <=> sP1), introduced(avatar_definition, [new_symbols(naming, [spl68_283])])).
fof(f2228, plain, (spl68_282 <=> sP2), introduced(avatar_definition, [new_symbols(naming, [spl68_282])])).
fof(f2221, plain, (spl68_281 <=> sP3), introduced(avatar_definition, [new_symbols(naming, [spl68_281])])).
fof(f2214, plain, (spl68_280 <=> sP4), introduced(avatar_definition, [new_symbols(naming, [spl68_280])])).
fof(f2207, plain, (spl68_279 <=> sP5), introduced(avatar_definition, [new_symbols(naming, [spl68_279])])).
fof(f2200, plain, (spl68_278 <=> sP6), introduced(avatar_definition, [new_symbols(naming, [spl68_278])])).
fof(f2193, plain, (spl68_277 <=> sP7), introduced(avatar_definition, [new_symbols(naming, [spl68_277])])).
fof(f2186, plain, (spl68_276 <=> sP8), introduced(avatar_definition, [new_symbols(naming, [spl68_276])])).
fof(f2179, plain, (spl68_275 <=> sP9), introduced(avatar_definition, [new_symbols(naming, [spl68_275])])).
fof(f2172, plain, (spl68_274 <=> sP10), introduced(avatar_definition, [new_symbols(naming, [spl68_274])])).
fof(f2165, plain, (spl68_273 <=> sP11), introduced(avatar_definition, [new_symbols(naming, [spl68_273])])).
fof(f2158, plain, (spl68_272 <=> sP12), introduced(avatar_definition, [new_symbols(naming, [spl68_272])])).
fof(f2151, plain, (spl68_271 <=> sP13), introduced(avatar_definition, [new_symbols(naming, [spl68_271])])).
fof(f2144, plain, (spl68_270 <=> sP14), introduced(avatar_definition, [new_symbols(naming, [spl68_270])])).
fof(f2137, plain, (spl68_269 <=> sP15), introduced(avatar_definition, [new_symbols(naming, [spl68_269])])).
fof(f2130, plain, (spl68_268 <=> sP16), introduced(avatar_definition, [new_symbols(naming, [spl68_268])])).
fof(f2123, plain, (spl68_267 <=> sP17), introduced(avatar_definition, [new_symbols(naming, [spl68_267])])).
fof(f2116, plain, (spl68_266 <=> sP18), introduced(avatar_definition, [new_symbols(naming, [spl68_266])])).
fof(f2109, plain, (spl68_265 <=> sP19), introduced(avatar_definition, [new_symbols(naming, [spl68_265])])).
fof(f2102, plain, (spl68_264 <=> sP20), introduced(avatar_definition, [new_symbols(naming, [spl68_264])])).
fof(f2095, plain, (spl68_263 <=> sP21), introduced(avatar_definition, [new_symbols(naming, [spl68_263])])).
fof(f2088, plain, (spl68_262 <=> sP22), introduced(avatar_definition, [new_symbols(naming, [spl68_262])])).
fof(f2081, plain, (spl68_261 <=> sP23), introduced(avatar_definition, [new_symbols(naming, [spl68_261])])).
fof(f652, plain, ((e14 = op1(e14, e14)) | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f48])).
fof(f48, plain, ((~ (e14 = op1(e14, e14)) & (e14 = op1(e14, e14)) & (e14 = op1(e14, e14))) | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(definition_folding, [], [f12, e47, e46, e45, e44, e43, e42, e41, e40, e39, e38, e37, e36, e35, e34, e33, e32, e31, e30, e29, e28, e27, e26, e25, e24])).
fof(f24, plain, ((~ (e10 = op1(e10, e10)) & (e10 = op1(e10, e10)) & (e10 = op1(e10, e10))) | ~ sP0), inference(usedef, [], [e24])).
fof(e24, plain, (sP0 <=> (~ (e10 = op1(e10, e10)) & (e10 = op1(e10, e10)) & (e10 = op1(e10, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f25, plain, ((~ (e10 = op1(e10, e11)) & (e10 = op1(e11, e11)) & (op1(e10, e10) = e11)) | ~ sP1), inference(usedef, [], [e25])).
fof(e25, plain, (sP1 <=> (~ (e10 = op1(e10, e11)) & (e10 = op1(e11, e11)) & (op1(e10, e10) = e11))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f26, plain, ((~ (e10 = op1(e10, e12)) & (e10 = op1(e12, e12)) & (op1(e10, e10) = e12)) | ~ sP2), inference(usedef, [], [e26])).
fof(e26, plain, (sP2 <=> (~ (e10 = op1(e10, e12)) & (e10 = op1(e12, e12)) & (op1(e10, e10) = e12))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f27, plain, ((~ (e10 = op1(e10, e13)) & (e10 = op1(e13, e13)) & (op1(e10, e10) = e13)) | ~ sP3), inference(usedef, [], [e27])).
fof(e27, plain, (sP3 <=> (~ (e10 = op1(e10, e13)) & (e10 = op1(e13, e13)) & (op1(e10, e10) = e13))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f28, plain, ((~ (e10 = op1(e10, e14)) & (e10 = op1(e14, e14)) & (op1(e10, e10) = e14)) | ~ sP4), inference(usedef, [], [e28])).
fof(e28, plain, (sP4 <=> (~ (e10 = op1(e10, e14)) & (e10 = op1(e14, e14)) & (op1(e10, e10) = e14))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f29, plain, ((~ (e11 = op1(e11, e10)) & (op1(e10, e10) = e11) & (e10 = op1(e11, e11))) | ~ sP5), inference(usedef, [], [e29])).
fof(e29, plain, (sP5 <=> (~ (e11 = op1(e11, e10)) & (op1(e10, e10) = e11) & (e10 = op1(e11, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f30, plain, ((~ (e11 = op1(e11, e11)) & (e11 = op1(e11, e11)) & (e11 = op1(e11, e11))) | ~ sP6), inference(usedef, [], [e30])).
fof(e30, plain, (sP6 <=> (~ (e11 = op1(e11, e11)) & (e11 = op1(e11, e11)) & (e11 = op1(e11, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f31, plain, ((~ (e11 = op1(e11, e12)) & (e11 = op1(e12, e12)) & (e12 = op1(e11, e11))) | ~ sP7), inference(usedef, [], [e31])).
fof(e31, plain, (sP7 <=> (~ (e11 = op1(e11, e12)) & (e11 = op1(e12, e12)) & (e12 = op1(e11, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f32, plain, ((~ (e11 = op1(e11, e13)) & (e11 = op1(e13, e13)) & (e13 = op1(e11, e11))) | ~ sP8), inference(usedef, [], [e32])).
fof(e32, plain, (sP8 <=> (~ (e11 = op1(e11, e13)) & (e11 = op1(e13, e13)) & (e13 = op1(e11, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f33, plain, ((~ (e11 = op1(e11, e14)) & (e11 = op1(e14, e14)) & (e14 = op1(e11, e11))) | ~ sP9), inference(usedef, [], [e33])).
fof(e33, plain, (sP9 <=> (~ (e11 = op1(e11, e14)) & (e11 = op1(e14, e14)) & (e14 = op1(e11, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f34, plain, ((~ (e12 = op1(e12, e10)) & (op1(e10, e10) = e12) & (e10 = op1(e12, e12))) | ~ sP10), inference(usedef, [], [e34])).
fof(e34, plain, (sP10 <=> (~ (e12 = op1(e12, e10)) & (op1(e10, e10) = e12) & (e10 = op1(e12, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f35, plain, ((~ (e12 = op1(e12, e11)) & (e12 = op1(e11, e11)) & (e11 = op1(e12, e12))) | ~ sP11), inference(usedef, [], [e35])).
fof(e35, plain, (sP11 <=> (~ (e12 = op1(e12, e11)) & (e12 = op1(e11, e11)) & (e11 = op1(e12, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f36, plain, ((~ (e12 = op1(e12, e12)) & (e12 = op1(e12, e12)) & (e12 = op1(e12, e12))) | ~ sP12), inference(usedef, [], [e36])).
fof(e36, plain, (sP12 <=> (~ (e12 = op1(e12, e12)) & (e12 = op1(e12, e12)) & (e12 = op1(e12, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f37, plain, ((~ (e12 = op1(e12, e13)) & (e12 = op1(e13, e13)) & (e13 = op1(e12, e12))) | ~ sP13), inference(usedef, [], [e37])).
fof(e37, plain, (sP13 <=> (~ (e12 = op1(e12, e13)) & (e12 = op1(e13, e13)) & (e13 = op1(e12, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f38, plain, ((~ (e12 = op1(e12, e14)) & (e12 = op1(e14, e14)) & (e14 = op1(e12, e12))) | ~ sP14), inference(usedef, [], [e38])).
fof(e38, plain, (sP14 <=> (~ (e12 = op1(e12, e14)) & (e12 = op1(e14, e14)) & (e14 = op1(e12, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f39, plain, ((~ (e13 = op1(e13, e10)) & (op1(e10, e10) = e13) & (e10 = op1(e13, e13))) | ~ sP15), inference(usedef, [], [e39])).
fof(e39, plain, (sP15 <=> (~ (e13 = op1(e13, e10)) & (op1(e10, e10) = e13) & (e10 = op1(e13, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f40, plain, ((~ (e13 = op1(e13, e11)) & (e13 = op1(e11, e11)) & (e11 = op1(e13, e13))) | ~ sP16), inference(usedef, [], [e40])).
fof(e40, plain, (sP16 <=> (~ (e13 = op1(e13, e11)) & (e13 = op1(e11, e11)) & (e11 = op1(e13, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f41, plain, ((~ (e13 = op1(e13, e12)) & (e13 = op1(e12, e12)) & (e12 = op1(e13, e13))) | ~ sP17), inference(usedef, [], [e41])).
fof(e41, plain, (sP17 <=> (~ (e13 = op1(e13, e12)) & (e13 = op1(e12, e12)) & (e12 = op1(e13, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f42, plain, ((~ (e13 = op1(e13, e13)) & (e13 = op1(e13, e13)) & (e13 = op1(e13, e13))) | ~ sP18), inference(usedef, [], [e42])).
fof(e42, plain, (sP18 <=> (~ (e13 = op1(e13, e13)) & (e13 = op1(e13, e13)) & (e13 = op1(e13, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f43, plain, ((~ (e13 = op1(e13, e14)) & (e13 = op1(e14, e14)) & (e14 = op1(e13, e13))) | ~ sP19), inference(usedef, [], [e43])).
fof(e43, plain, (sP19 <=> (~ (e13 = op1(e13, e14)) & (e13 = op1(e14, e14)) & (e14 = op1(e13, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP19])])).
fof(f44, plain, ((~ (e14 = op1(e14, e10)) & (op1(e10, e10) = e14) & (e10 = op1(e14, e14))) | ~ sP20), inference(usedef, [], [e44])).
fof(e44, plain, (sP20 <=> (~ (e14 = op1(e14, e10)) & (op1(e10, e10) = e14) & (e10 = op1(e14, e14)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP20])])).
fof(f45, plain, ((~ (e14 = op1(e14, e11)) & (e14 = op1(e11, e11)) & (e11 = op1(e14, e14))) | ~ sP21), inference(usedef, [], [e45])).
fof(e45, plain, (sP21 <=> (~ (e14 = op1(e14, e11)) & (e14 = op1(e11, e11)) & (e11 = op1(e14, e14)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP21])])).
fof(f46, plain, ((~ (e14 = op1(e14, e12)) & (e14 = op1(e12, e12)) & (e12 = op1(e14, e14))) | ~ sP22), inference(usedef, [], [e46])).
fof(e46, plain, (sP22 <=> (~ (e14 = op1(e14, e12)) & (e14 = op1(e12, e12)) & (e12 = op1(e14, e14)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP22])])).
fof(f47, plain, ((~ (e14 = op1(e14, e13)) & (e14 = op1(e13, e13)) & (e13 = op1(e14, e14))) | ~ sP23), inference(usedef, [], [e47])).
fof(e47, plain, (sP23 <=> (~ (e14 = op1(e14, e13)) & (e14 = op1(e13, e13)) & (e13 = op1(e14, e14)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP23])])).
fof(f12, plain, ((~ (e14 = op1(e14, e14)) & (e14 = op1(e14, e14)) & (e14 = op1(e14, e14))) | (~ (e14 = op1(e14, e13)) & (e14 = op1(e13, e13)) & (e13 = op1(e14, e14))) | (~ (e14 = op1(e14, e12)) & (e14 = op1(e12, e12)) & (e12 = op1(e14, e14))) | (~ (e14 = op1(e14, e11)) & (e14 = op1(e11, e11)) & (e11 = op1(e14, e14))) | (~ (e14 = op1(e14, e10)) & (op1(e10, e10) = e14) & (e10 = op1(e14, e14))) | (~ (e13 = op1(e13, e14)) & (e13 = op1(e14, e14)) & (e14 = op1(e13, e13))) | (~ (e13 = op1(e13, e13)) & (e13 = op1(e13, e13)) & (e13 = op1(e13, e13))) | (~ (e13 = op1(e13, e12)) & (e13 = op1(e12, e12)) & (e12 = op1(e13, e13))) | (~ (e13 = op1(e13, e11)) & (e13 = op1(e11, e11)) & (e11 = op1(e13, e13))) | (~ (e13 = op1(e13, e10)) & (op1(e10, e10) = e13) & (e10 = op1(e13, e13))) | (~ (e12 = op1(e12, e14)) & (e12 = op1(e14, e14)) & (e14 = op1(e12, e12))) | (~ (e12 = op1(e12, e13)) & (e12 = op1(e13, e13)) & (e13 = op1(e12, e12))) | (~ (e12 = op1(e12, e12)) & (e12 = op1(e12, e12)) & (e12 = op1(e12, e12))) | (~ (e12 = op1(e12, e11)) & (e12 = op1(e11, e11)) & (e11 = op1(e12, e12))) | (~ (e12 = op1(e12, e10)) & (op1(e10, e10) = e12) & (e10 = op1(e12, e12))) | (~ (e11 = op1(e11, e14)) & (e11 = op1(e14, e14)) & (e14 = op1(e11, e11))) | (~ (e11 = op1(e11, e13)) & (e11 = op1(e13, e13)) & (e13 = op1(e11, e11))) | (~ (e11 = op1(e11, e12)) & (e11 = op1(e12, e12)) & (e12 = op1(e11, e11))) | (~ (e11 = op1(e11, e11)) & (e11 = op1(e11, e11)) & (e11 = op1(e11, e11))) | (~ (e11 = op1(e11, e10)) & (op1(e10, e10) = e11) & (e10 = op1(e11, e11))) | (~ (e10 = op1(e10, e14)) & (e10 = op1(e14, e14)) & (op1(e10, e10) = e14)) | (~ (e10 = op1(e10, e13)) & (e10 = op1(e13, e13)) & (op1(e10, e10) = e13)) | (~ (e10 = op1(e10, e12)) & (e10 = op1(e12, e12)) & (op1(e10, e10) = e12)) | (~ (e10 = op1(e10, e11)) & (e10 = op1(e11, e11)) & (op1(e10, e10) = e11)) | (~ (e10 = op1(e10, e10)) & (e10 = op1(e10, e10)) & (e10 = op1(e10, e10)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG052+1.p', ax12)).
fof(f2248, plain, (spl68_284 | spl68_283 | spl68_282 | spl68_281 | spl68_280 | spl68_279 | spl68_278 | spl68_277 | spl68_276 | spl68_275 | spl68_274 | spl68_273 | spl68_272 | spl68_271 | spl68_270 | spl68_269 | spl68_268 | spl68_267 | spl68_266 | spl68_265 | spl68_264 | spl68_263 | spl68_262 | spl68_261 | ~ spl68_5), inference(avatar_split_clause, [], [f654, f905, f2081, f2088, f2095, f2102, f2109, f2116, f2123, f2130, f2137, f2144, f2151, f2158, f2165, f2172, f2179, f2186, f2193, f2200, f2207, f2214, f2221, f2228, f2235, f2242])).
fof(f654, plain, (~ (e14 = op1(e14, e14)) | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f48])).
fof(f2247, plain, (~ spl68_284 | spl68_121), inference(avatar_split_clause, [], [f649, f1393, f2242])).
fof(f649, plain, ((e10 = op1(e10, e10)) | ~ sP0), inference(cnf_transformation, [], [f118])).
fof(f118, plain, ((~ (e10 = op1(e10, e10)) & (e10 = op1(e10, e10)) & (e10 = op1(e10, e10))) | ~ sP0), inference(nnf_transformation, [], [f24])).
fof(f2245, plain, (~ spl68_284 | ~ spl68_121), inference(avatar_split_clause, [], [f651, f1393, f2242])).
fof(f651, plain, (~ (e10 = op1(e10, e10)) | ~ sP0), inference(cnf_transformation, [], [f118])).
fof(f2239, plain, (~ spl68_283 | spl68_91), inference(avatar_split_clause, [], [f647, f1267, f2235])).
fof(f647, plain, ((e10 = op1(e11, e11)) | ~ sP1), inference(cnf_transformation, [], [f117])).
fof(f117, plain, ((~ (e10 = op1(e10, e11)) & (e10 = op1(e11, e11)) & (op1(e10, e10) = e11)) | ~ sP1), inference(nnf_transformation, [], [f25])).
fof(f2233, plain, (~ spl68_282 | spl68_123), inference(avatar_split_clause, [], [f643, f1401, f2228])).
fof(f643, plain, ((op1(e10, e10) = e12) | ~ sP2), inference(cnf_transformation, [], [f116])).
fof(f116, plain, ((~ (e10 = op1(e10, e12)) & (e10 = op1(e12, e12)) & (op1(e10, e10) = e12)) | ~ sP2), inference(nnf_transformation, [], [f26])).
fof(f2231, plain, (~ spl68_282 | ~ spl68_111), inference(avatar_split_clause, [], [f645, f1351, f2228])).
fof(f645, plain, (~ (e10 = op1(e10, e12)) | ~ sP2), inference(cnf_transformation, [], [f116])).
fof(f2226, plain, (~ spl68_281 | spl68_124), inference(avatar_split_clause, [], [f640, f1405, f2221])).
fof(f640, plain, ((op1(e10, e10) = e13) | ~ sP3), inference(cnf_transformation, [], [f115])).
fof(f115, plain, ((~ (e10 = op1(e10, e13)) & (e10 = op1(e13, e13)) & (op1(e10, e10) = e13)) | ~ sP3), inference(nnf_transformation, [], [f27])).
fof(f2225, plain, (~ spl68_281 | spl68_31), inference(avatar_split_clause, [], [f641, f1015, f2221])).
fof(f641, plain, ((e10 = op1(e13, e13)) | ~ sP3), inference(cnf_transformation, [], [f115])).
fof(f2218, plain, (~ spl68_280 | spl68_1), inference(avatar_split_clause, [], [f638, f889, f2214])).
fof(f638, plain, ((e10 = op1(e14, e14)) | ~ sP4), inference(cnf_transformation, [], [f114])).
fof(f114, plain, ((~ (e10 = op1(e10, e14)) & (e10 = op1(e14, e14)) & (op1(e10, e10) = e14)) | ~ sP4), inference(nnf_transformation, [], [f28])).
fof(f2212, plain, (~ spl68_279 | spl68_91), inference(avatar_split_clause, [], [f634, f1267, f2207])).
fof(f634, plain, ((e10 = op1(e11, e11)) | ~ sP5), inference(cnf_transformation, [], [f113])).
fof(f113, plain, ((~ (e11 = op1(e11, e10)) & (op1(e10, e10) = e11) & (e10 = op1(e11, e11))) | ~ sP5), inference(nnf_transformation, [], [f29])).
fof(f2205, plain, (~ spl68_278 | spl68_92), inference(avatar_split_clause, [], [f631, f1271, f2200])).
fof(f1271, plain, (spl68_92 <=> (e11 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl68_92])])).
fof(f631, plain, ((e11 = op1(e11, e11)) | ~ sP6), inference(cnf_transformation, [], [f112])).
fof(f112, plain, ((~ (e11 = op1(e11, e11)) & (e11 = op1(e11, e11)) & (e11 = op1(e11, e11))) | ~ sP6), inference(nnf_transformation, [], [f30])).
fof(f2203, plain, (~ spl68_278 | ~ spl68_92), inference(avatar_split_clause, [], [f633, f1271, f2200])).
fof(f633, plain, (~ (e11 = op1(e11, e11)) | ~ sP6), inference(cnf_transformation, [], [f112])).
fof(f2197, plain, (~ spl68_277 | spl68_62), inference(avatar_split_clause, [], [f629, f1145, f2193])).
fof(f629, plain, ((e11 = op1(e12, e12)) | ~ sP7), inference(cnf_transformation, [], [f111])).
fof(f111, plain, ((~ (e11 = op1(e11, e12)) & (e11 = op1(e12, e12)) & (e12 = op1(e11, e11))) | ~ sP7), inference(nnf_transformation, [], [f31])).
fof(f2196, plain, (~ spl68_277 | ~ spl68_87), inference(avatar_split_clause, [], [f630, f1250, f2193])).
fof(f630, plain, (~ (e11 = op1(e11, e12)) | ~ sP7), inference(cnf_transformation, [], [f111])).
fof(f2191, plain, (~ spl68_276 | spl68_94), inference(avatar_split_clause, [], [f625, f1279, f2186])).
fof(f625, plain, ((e13 = op1(e11, e11)) | ~ sP8), inference(cnf_transformation, [], [f110])).
fof(f110, plain, ((~ (e11 = op1(e11, e13)) & (e11 = op1(e13, e13)) & (e13 = op1(e11, e11))) | ~ sP8), inference(nnf_transformation, [], [f32])).
fof(f2184, plain, (~ spl68_275 | spl68_95), inference(avatar_split_clause, [], [f622, f1283, f2179])).
fof(f622, plain, ((e14 = op1(e11, e11)) | ~ sP9), inference(cnf_transformation, [], [f109])).
fof(f109, plain, ((~ (e11 = op1(e11, e14)) & (e11 = op1(e14, e14)) & (e14 = op1(e11, e11))) | ~ sP9), inference(nnf_transformation, [], [f33])).
fof(f2177, plain, (~ spl68_274 | spl68_61), inference(avatar_split_clause, [], [f619, f1141, f2172])).
fof(f619, plain, ((e10 = op1(e12, e12)) | ~ sP10), inference(cnf_transformation, [], [f108])).
fof(f108, plain, ((~ (e12 = op1(e12, e10)) & (op1(e10, e10) = e12) & (e10 = op1(e12, e12))) | ~ sP10), inference(nnf_transformation, [], [f34])).
fof(f2175, plain, (~ spl68_274 | ~ spl68_73), inference(avatar_split_clause, [], [f621, f1191, f2172])).
fof(f621, plain, (~ (e12 = op1(e12, e10)) | ~ sP10), inference(cnf_transformation, [], [f108])).
fof(f2170, plain, (~ spl68_273 | spl68_62), inference(avatar_split_clause, [], [f616, f1145, f2165])).
fof(f616, plain, ((e11 = op1(e12, e12)) | ~ sP11), inference(cnf_transformation, [], [f107])).
fof(f107, plain, ((~ (e12 = op1(e12, e11)) & (e12 = op1(e11, e11)) & (e11 = op1(e12, e12))) | ~ sP11), inference(nnf_transformation, [], [f35])).
fof(f2163, plain, (~ spl68_272 | spl68_63), inference(avatar_split_clause, [], [f613, f1149, f2158])).
fof(f613, plain, ((e12 = op1(e12, e12)) | ~ sP12), inference(cnf_transformation, [], [f106])).
fof(f106, plain, ((~ (e12 = op1(e12, e12)) & (e12 = op1(e12, e12)) & (e12 = op1(e12, e12))) | ~ sP12), inference(nnf_transformation, [], [f36])).
fof(f2161, plain, (~ spl68_272 | ~ spl68_63), inference(avatar_split_clause, [], [f615, f1149, f2158])).
fof(f615, plain, (~ (e12 = op1(e12, e12)) | ~ sP12), inference(cnf_transformation, [], [f106])).
fof(f2155, plain, (~ spl68_271 | spl68_33), inference(avatar_split_clause, [], [f611, f1023, f2151])).
fof(f611, plain, ((e12 = op1(e13, e13)) | ~ sP13), inference(cnf_transformation, [], [f105])).
fof(f105, plain, ((~ (e12 = op1(e12, e13)) & (e12 = op1(e13, e13)) & (e13 = op1(e12, e12))) | ~ sP13), inference(nnf_transformation, [], [f37])).
fof(f2148, plain, (~ spl68_270 | spl68_3), inference(avatar_split_clause, [], [f608, f897, f2144])).
fof(f608, plain, ((e12 = op1(e14, e14)) | ~ sP14), inference(cnf_transformation, [], [f104])).
fof(f104, plain, ((~ (e12 = op1(e12, e14)) & (e12 = op1(e14, e14)) & (e14 = op1(e12, e12))) | ~ sP14), inference(nnf_transformation, [], [f38])).
fof(f2142, plain, (~ spl68_269 | spl68_31), inference(avatar_split_clause, [], [f604, f1015, f2137])).
fof(f604, plain, ((e10 = op1(e13, e13)) | ~ sP15), inference(cnf_transformation, [], [f103])).
fof(f103, plain, ((~ (e13 = op1(e13, e10)) & (op1(e10, e10) = e13) & (e10 = op1(e13, e13))) | ~ sP15), inference(nnf_transformation, [], [f39])).
fof(f2141, plain, (~ spl68_269 | spl68_124), inference(avatar_split_clause, [], [f605, f1405, f2137])).
fof(f605, plain, ((op1(e10, e10) = e13) | ~ sP15), inference(cnf_transformation, [], [f103])).
fof(f2140, plain, (~ spl68_269 | ~ spl68_49), inference(avatar_split_clause, [], [f606, f1090, f2137])).
fof(f606, plain, (~ (e13 = op1(e13, e10)) | ~ sP15), inference(cnf_transformation, [], [f103])).
fof(f2134, plain, (~ spl68_268 | spl68_94), inference(avatar_split_clause, [], [f602, f1279, f2130])).
fof(f602, plain, ((e13 = op1(e11, e11)) | ~ sP16), inference(cnf_transformation, [], [f102])).
fof(f102, plain, ((~ (e13 = op1(e13, e11)) & (e13 = op1(e11, e11)) & (e11 = op1(e13, e13))) | ~ sP16), inference(nnf_transformation, [], [f40])).
fof(f2128, plain, (~ spl68_267 | spl68_33), inference(avatar_split_clause, [], [f598, f1023, f2123])).
fof(f598, plain, ((e12 = op1(e13, e13)) | ~ sP17), inference(cnf_transformation, [], [f101])).
fof(f101, plain, ((~ (e13 = op1(e13, e12)) & (e13 = op1(e12, e12)) & (e12 = op1(e13, e13))) | ~ sP17), inference(nnf_transformation, [], [f41])).
fof(f2121, plain, (~ spl68_266 | spl68_34), inference(avatar_split_clause, [], [f595, f1027, f2116])).
fof(f595, plain, ((e13 = op1(e13, e13)) | ~ sP18), inference(cnf_transformation, [], [f100])).
fof(f100, plain, ((~ (e13 = op1(e13, e13)) & (e13 = op1(e13, e13)) & (e13 = op1(e13, e13))) | ~ sP18), inference(nnf_transformation, [], [f42])).
fof(f2113, plain, (~ spl68_265 | spl68_4), inference(avatar_split_clause, [], [f593, f901, f2109])).
fof(f593, plain, ((e13 = op1(e14, e14)) | ~ sP19), inference(cnf_transformation, [], [f99])).
fof(f99, plain, ((~ (e13 = op1(e13, e14)) & (e13 = op1(e14, e14)) & (e14 = op1(e13, e13))) | ~ sP19), inference(nnf_transformation, [], [f43])).
fof(f2107, plain, (~ spl68_264 | spl68_1), inference(avatar_split_clause, [], [f589, f889, f2102])).
fof(f589, plain, ((e10 = op1(e14, e14)) | ~ sP20), inference(cnf_transformation, [], [f98])).
fof(f98, plain, ((~ (e14 = op1(e14, e10)) & (op1(e10, e10) = e14) & (e10 = op1(e14, e14))) | ~ sP20), inference(nnf_transformation, [], [f44])).
fof(f2099, plain, (~ spl68_263 | spl68_95), inference(avatar_split_clause, [], [f587, f1283, f2095])).
fof(f587, plain, ((e14 = op1(e11, e11)) | ~ sP21), inference(cnf_transformation, [], [f97])).
fof(f97, plain, ((~ (e14 = op1(e14, e11)) & (e14 = op1(e11, e11)) & (e11 = op1(e14, e14))) | ~ sP21), inference(nnf_transformation, [], [f45])).
fof(f2093, plain, (~ spl68_262 | spl68_3), inference(avatar_split_clause, [], [f583, f897, f2088])).
fof(f583, plain, ((e12 = op1(e14, e14)) | ~ sP22), inference(cnf_transformation, [], [f96])).
fof(f96, plain, ((~ (e14 = op1(e14, e12)) & (e14 = op1(e12, e12)) & (e12 = op1(e14, e14))) | ~ sP22), inference(nnf_transformation, [], [f46])).
fof(f2086, plain, (~ spl68_261 | spl68_4), inference(avatar_split_clause, [], [f580, f901, f2081])).
fof(f580, plain, ((e13 = op1(e14, e14)) | ~ sP23), inference(cnf_transformation, [], [f95])).
fof(f95, plain, ((~ (e14 = op1(e14, e13)) & (e14 = op1(e13, e13)) & (e13 = op1(e14, e14))) | ~ sP23), inference(nnf_transformation, [], [f47])).
fof(f2076, plain, (spl68_252 | spl68_227 | spl68_202 | spl68_177 | spl68_152), inference(avatar_split_clause, [], [f288, f1573, f1678, f1783, f1888, f1993])).
fof(f288, plain, ((e21 = op2(e24, e20)) | (e21 = op2(e23, e20)) | (e21 = op2(e22, e20)) | (e21 = op2(e21, e20)) | (op2(e20, e20) = e21)), inference(cnf_transformation, [], [f6])).
fof(f6, plain, (((e24 = op2(e24, e24)) | (e24 = op2(e23, e24)) | (e24 = op2(e22, e24)) | (e24 = op2(e21, e24)) | (e24 = op2(e20, e24))) & ((e24 = op2(e24, e24)) | (e24 = op2(e24, e23)) | (e24 = op2(e24, e22)) | (e24 = op2(e24, e21)) | (e24 = op2(e24, e20))) & ((e23 = op2(e24, e24)) | (e23 = op2(e23, e24)) | (e23 = op2(e22, e24)) | (e23 = op2(e21, e24)) | (e23 = op2(e20, e24))) & ((e23 = op2(e24, e24)) | (e23 = op2(e24, e23)) | (e23 = op2(e24, e22)) | (e23 = op2(e24, e21)) | (e23 = op2(e24, e20))) & ((e22 = op2(e24, e24)) | (e22 = op2(e23, e24)) | (e22 = op2(e22, e24)) | (e22 = op2(e21, e24)) | (e22 = op2(e20, e24))) & ((e22 = op2(e24, e24)) | (e22 = op2(e24, e23)) | (e22 = op2(e24, e22)) | (e22 = op2(e24, e21)) | (e22 = op2(e24, e20))) & ((e21 = op2(e24, e24)) | (e21 = op2(e23, e24)) | (e21 = op2(e22, e24)) | (e21 = op2(e21, e24)) | (e21 = op2(e20, e24))) & ((e21 = op2(e24, e24)) | (e21 = op2(e24, e23)) | (e21 = op2(e24, e22)) | (e21 = op2(e24, e21)) | (e21 = op2(e24, e20))) & ((e20 = op2(e24, e24)) | (e20 = op2(e23, e24)) | (e20 = op2(e22, e24)) | (e20 = op2(e21, e24)) | (e20 = op2(e20, e24))) & ((e20 = op2(e24, e24)) | (e20 = op2(e24, e23)) | (e20 = op2(e24, e22)) | (e20 = op2(e24, e21)) | (e20 = op2(e24, e20))) & ((e24 = op2(e24, e23)) | (e24 = op2(e23, e23)) | (e24 = op2(e22, e23)) | (e24 = op2(e21, e23)) | (e24 = op2(e20, e23))) & ((e24 = op2(e23, e24)) | (e24 = op2(e23, e23)) | (e24 = op2(e23, e22)) | (e24 = op2(e23, e21)) | (e24 = op2(e23, e20))) & ((e23 = op2(e24, e23)) | (e23 = op2(e23, e23)) | (e23 = op2(e22, e23)) | (e23 = op2(e21, e23)) | (e23 = op2(e20, e23))) & ((e23 = op2(e23, e24)) | (e23 = op2(e23, e23)) | (e23 = op2(e23, e22)) | (e23 = op2(e23, e21)) | (e23 = op2(e23, e20))) & ((e22 = op2(e24, e23)) | (e22 = op2(e23, e23)) | (e22 = op2(e22, e23)) | (e22 = op2(e21, e23)) | (e22 = op2(e20, e23))) & ((e22 = op2(e23, e24)) | (e22 = op2(e23, e23)) | (e22 = op2(e23, e22)) | (e22 = op2(e23, e21)) | (e22 = op2(e23, e20))) & ((e21 = op2(e24, e23)) | (e21 = op2(e23, e23)) | (e21 = op2(e22, e23)) | (e21 = op2(e21, e23)) | (e21 = op2(e20, e23))) & ((e21 = op2(e23, e24)) | (e21 = op2(e23, e23)) | (e21 = op2(e23, e22)) | (e21 = op2(e23, e21)) | (e21 = op2(e23, e20))) & ((e20 = op2(e24, e23)) | (e20 = op2(e23, e23)) | (e20 = op2(e22, e23)) | (e20 = op2(e21, e23)) | (e20 = op2(e20, e23))) & ((e20 = op2(e23, e24)) | (e20 = op2(e23, e23)) | (e20 = op2(e23, e22)) | (e20 = op2(e23, e21)) | (e20 = op2(e23, e20))) & ((e24 = op2(e24, e22)) | (e24 = op2(e23, e22)) | (e24 = op2(e22, e22)) | (e24 = op2(e21, e22)) | (e24 = op2(e20, e22))) & ((e24 = op2(e22, e24)) | (e24 = op2(e22, e23)) | (e24 = op2(e22, e22)) | (e24 = op2(e22, e21)) | (e24 = op2(e22, e20))) & ((e23 = op2(e24, e22)) | (e23 = op2(e23, e22)) | (e23 = op2(e22, e22)) | (e23 = op2(e21, e22)) | (e23 = op2(e20, e22))) & ((e23 = op2(e22, e24)) | (e23 = op2(e22, e23)) | (e23 = op2(e22, e22)) | (e23 = op2(e22, e21)) | (e23 = op2(e22, e20))) & ((e22 = op2(e24, e22)) | (e22 = op2(e23, e22)) | (e22 = op2(e22, e22)) | (e22 = op2(e21, e22)) | (e22 = op2(e20, e22))) & ((e22 = op2(e22, e24)) | (e22 = op2(e22, e23)) | (e22 = op2(e22, e22)) | (e22 = op2(e22, e21)) | (e22 = op2(e22, e20))) & ((e21 = op2(e24, e22)) | (e21 = op2(e23, e22)) | (e21 = op2(e22, e22)) | (e21 = op2(e21, e22)) | (e21 = op2(e20, e22))) & ((e21 = op2(e22, e24)) | (e21 = op2(e22, e23)) | (e21 = op2(e22, e22)) | (e21 = op2(e22, e21)) | (e21 = op2(e22, e20))) & ((e20 = op2(e24, e22)) | (e20 = op2(e23, e22)) | (e20 = op2(e22, e22)) | (e20 = op2(e21, e22)) | (e20 = op2(e20, e22))) & ((e20 = op2(e22, e24)) | (e20 = op2(e22, e23)) | (e20 = op2(e22, e22)) | (e20 = op2(e22, e21)) | (e20 = op2(e22, e20))) & ((e24 = op2(e24, e21)) | (e24 = op2(e23, e21)) | (e24 = op2(e22, e21)) | (e24 = op2(e21, e21)) | (e24 = op2(e20, e21))) & ((e24 = op2(e21, e24)) | (e24 = op2(e21, e23)) | (e24 = op2(e21, e22)) | (e24 = op2(e21, e21)) | (e24 = op2(e21, e20))) & ((e23 = op2(e24, e21)) | (e23 = op2(e23, e21)) | (e23 = op2(e22, e21)) | (e23 = op2(e21, e21)) | (e23 = op2(e20, e21))) & ((e23 = op2(e21, e24)) | (e23 = op2(e21, e23)) | (e23 = op2(e21, e22)) | (e23 = op2(e21, e21)) | (e23 = op2(e21, e20))) & ((e22 = op2(e24, e21)) | (e22 = op2(e23, e21)) | (e22 = op2(e22, e21)) | (e22 = op2(e21, e21)) | (e22 = op2(e20, e21))) & ((e22 = op2(e21, e24)) | (e22 = op2(e21, e23)) | (e22 = op2(e21, e22)) | (e22 = op2(e21, e21)) | (e22 = op2(e21, e20))) & ((e21 = op2(e24, e21)) | (e21 = op2(e23, e21)) | (e21 = op2(e22, e21)) | (e21 = op2(e21, e21)) | (e21 = op2(e20, e21))) & ((e21 = op2(e21, e24)) | (e21 = op2(e21, e23)) | (e21 = op2(e21, e22)) | (e21 = op2(e21, e21)) | (e21 = op2(e21, e20))) & ((e20 = op2(e24, e21)) | (e20 = op2(e23, e21)) | (e20 = op2(e22, e21)) | (e20 = op2(e21, e21)) | (e20 = op2(e20, e21))) & ((e20 = op2(e21, e24)) | (e20 = op2(e21, e23)) | (e20 = op2(e21, e22)) | (e20 = op2(e21, e21)) | (e20 = op2(e21, e20))) & ((e24 = op2(e24, e20)) | (e24 = op2(e23, e20)) | (e24 = op2(e22, e20)) | (e24 = op2(e21, e20)) | (op2(e20, e20) = e24)) & ((e24 = op2(e20, e24)) | (e24 = op2(e20, e23)) | (e24 = op2(e20, e22)) | (e24 = op2(e20, e21)) | (op2(e20, e20) = e24)) & ((e23 = op2(e24, e20)) | (e23 = op2(e23, e20)) | (e23 = op2(e22, e20)) | (e23 = op2(e21, e20)) | (op2(e20, e20) = e23)) & ((e23 = op2(e20, e24)) | (e23 = op2(e20, e23)) | (e23 = op2(e20, e22)) | (e23 = op2(e20, e21)) | (op2(e20, e20) = e23)) & ((e22 = op2(e24, e20)) | (e22 = op2(e23, e20)) | (e22 = op2(e22, e20)) | (e22 = op2(e21, e20)) | (op2(e20, e20) = e22)) & ((e22 = op2(e20, e24)) | (e22 = op2(e20, e23)) | (e22 = op2(e20, e22)) | (e22 = op2(e20, e21)) | (op2(e20, e20) = e22)) & ((e21 = op2(e24, e20)) | (e21 = op2(e23, e20)) | (e21 = op2(e22, e20)) | (e21 = op2(e21, e20)) | (op2(e20, e20) = e21)) & ((e21 = op2(e20, e24)) | (e21 = op2(e20, e23)) | (e21 = op2(e20, e22)) | (e21 = op2(e20, e21)) | (op2(e20, e20) = e21)) & ((e20 = op2(e24, e20)) | (e20 = op2(e23, e20)) | (e20 = op2(e22, e20)) | (e20 = op2(e21, e20)) | (e20 = op2(e20, e20))) & ((e20 = op2(e20, e24)) | (e20 = op2(e20, e23)) | (e20 = op2(e20, e22)) | (e20 = op2(e20, e21)) | (e20 = op2(e20, e20)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG052+1.p', ax6)).
fof(f2072, plain, (spl68_254 | spl68_229 | spl68_204 | spl68_179 | spl68_154), inference(avatar_split_clause, [], [f292, f1581, f1686, f1791, f1896, f2001])).
fof(f292, plain, ((e23 = op2(e24, e20)) | (e23 = op2(e23, e20)) | (e23 = op2(e22, e20)) | (e23 = op2(e21, e20)) | (op2(e20, e20) = e23)), inference(cnf_transformation, [], [f6])).
fof(f2071, plain, (spl68_255 | spl68_250 | spl68_245 | spl68_240 | spl68_235), inference(avatar_split_clause, [], [f293, f1921, f1942, f1963, f1984, f2005])).
fof(f293, plain, ((e24 = op2(e20, e24)) | (e24 = op2(e20, e23)) | (e24 = op2(e20, e22)) | (e24 = op2(e20, e21)) | (op2(e20, e20) = e24)), inference(cnf_transformation, [], [f6])).
fof(f2068, plain, (spl68_246 | spl68_221 | spl68_196 | spl68_171 | spl68_146), inference(avatar_split_clause, [], [f296, f1548, f1653, f1758, f1863, f1968])).
fof(f296, plain, ((e20 = op2(e24, e21)) | (e20 = op2(e23, e21)) | (e20 = op2(e22, e21)) | (e20 = op2(e21, e21)) | (e20 = op2(e20, e21))), inference(cnf_transformation, [], [f6])).
fof(f2065, plain, (spl68_228 | spl68_223 | spl68_218 | spl68_213 | spl68_208), inference(avatar_split_clause, [], [f299, f1808, f1829, f1850, f1871, f1892])).
fof(f299, plain, ((e22 = op2(e21, e24)) | (e22 = op2(e21, e23)) | (e22 = op2(e21, e22)) | (e22 = op2(e21, e21)) | (e22 = op2(e21, e20))), inference(cnf_transformation, [], [f6])).
fof(f2064, plain, (spl68_248 | spl68_223 | spl68_198 | spl68_173 | spl68_148), inference(avatar_split_clause, [], [f300, f1556, f1661, f1766, f1871, f1976])).
fof(f300, plain, ((e22 = op2(e24, e21)) | (e22 = op2(e23, e21)) | (e22 = op2(e22, e21)) | (e22 = op2(e21, e21)) | (e22 = op2(e20, e21))), inference(cnf_transformation, [], [f6])).
fof(f2063, plain, (spl68_229 | spl68_224 | spl68_219 | spl68_214 | spl68_209), inference(avatar_split_clause, [], [f301, f1812, f1833, f1854, f1875, f1896])).
fof(f301, plain, ((e23 = op2(e21, e24)) | (e23 = op2(e21, e23)) | (e23 = op2(e21, e22)) | (e23 = op2(e21, e21)) | (e23 = op2(e21, e20))), inference(cnf_transformation, [], [f6])).
fof(f2060, plain, (spl68_250 | spl68_225 | spl68_200 | spl68_175 | spl68_150), inference(avatar_split_clause, [], [f304, f1564, f1669, f1774, f1879, f1984])).
fof(f304, plain, ((e24 = op2(e24, e21)) | (e24 = op2(e23, e21)) | (e24 = op2(e22, e21)) | (e24 = op2(e21, e21)) | (e24 = op2(e20, e21))), inference(cnf_transformation, [], [f6])).
fof(f2059, plain, (spl68_201 | spl68_196 | spl68_191 | spl68_186 | spl68_181), inference(avatar_split_clause, [], [f305, f1695, f1716, f1737, f1758, f1779])).
fof(f305, plain, ((e20 = op2(e22, e24)) | (e20 = op2(e22, e23)) | (e20 = op2(e22, e22)) | (e20 = op2(e22, e21)) | (e20 = op2(e22, e20))), inference(cnf_transformation, [], [f6])).
fof(f2056, plain, (spl68_242 | spl68_217 | spl68_192 | spl68_167 | spl68_142), inference(avatar_split_clause, [], [f308, f1531, f1636, f1741, f1846, f1951])).
fof(f308, plain, ((e21 = op2(e24, e22)) | (e21 = op2(e23, e22)) | (e21 = op2(e22, e22)) | (e21 = op2(e21, e22)) | (e21 = op2(e20, e22))), inference(cnf_transformation, [], [f6])).
fof(f2047, plain, (spl68_177 | spl68_172 | spl68_167 | spl68_162 | spl68_157), inference(avatar_split_clause, [], [f317, f1594, f1615, f1636, f1657, f1678])).
fof(f317, plain, ((e21 = op2(e23, e24)) | (e21 = op2(e23, e23)) | (e21 = op2(e23, e22)) | (e21 = op2(e23, e21)) | (e21 = op2(e23, e20))), inference(cnf_transformation, [], [f6])).
fof(f2046, plain, (spl68_237 | spl68_212 | spl68_187 | spl68_162 | spl68_137), inference(avatar_split_clause, [], [f318, f1510, f1615, f1720, f1825, f1930])).
fof(f318, plain, ((e21 = op2(e24, e23)) | (e21 = op2(e23, e23)) | (e21 = op2(e22, e23)) | (e21 = op2(e21, e23)) | (e21 = op2(e20, e23))), inference(cnf_transformation, [], [f6])).
fof(f2041, plain, (spl68_180 | spl68_175 | spl68_170 | spl68_165 | spl68_160), inference(avatar_split_clause, [], [f323, f1606, f1627, f1648, f1669, f1690])).
fof(f323, plain, ((e24 = op2(e23, e24)) | (e24 = op2(e23, e23)) | (e24 = op2(e23, e22)) | (e24 = op2(e23, e21)) | (e24 = op2(e23, e20))), inference(cnf_transformation, [], [f6])).
fof(f2040, plain, (spl68_240 | spl68_215 | spl68_190 | spl68_165 | spl68_140), inference(avatar_split_clause, [], [f324, f1522, f1627, f1732, f1837, f1942])).
fof(f324, plain, ((e24 = op2(e24, e23)) | (e24 = op2(e23, e23)) | (e24 = op2(e22, e23)) | (e24 = op2(e21, e23)) | (e24 = op2(e20, e23))), inference(cnf_transformation, [], [f6])).
fof(f2039, plain, (spl68_151 | spl68_146 | spl68_141 | spl68_136 | spl68_131), inference(avatar_split_clause, [], [f325, f1485, f1506, f1527, f1548, f1569])).
fof(f325, plain, ((e20 = op2(e24, e24)) | (e20 = op2(e24, e23)) | (e20 = op2(e24, e22)) | (e20 = op2(e24, e21)) | (e20 = op2(e24, e20))), inference(cnf_transformation, [], [f6])).
fof(f2035, plain, (spl68_153 | spl68_148 | spl68_143 | spl68_138 | spl68_133), inference(avatar_split_clause, [], [f329, f1493, f1514, f1535, f1556, f1577])).
fof(f329, plain, ((e22 = op2(e24, e24)) | (e22 = op2(e24, e23)) | (e22 = op2(e24, e22)) | (e22 = op2(e24, e21)) | (e22 = op2(e24, e20))), inference(cnf_transformation, [], [f6])).
fof(f2032, plain, (spl68_234 | spl68_209 | spl68_184 | spl68_159 | spl68_134), inference(avatar_split_clause, [], [f332, f1497, f1602, f1707, f1812, f1917])).
fof(f332, plain, ((e23 = op2(e24, e24)) | (e23 = op2(e23, e24)) | (e23 = op2(e22, e24)) | (e23 = op2(e21, e24)) | (e23 = op2(e20, e24))), inference(cnf_transformation, [], [f6])).
fof(f2029, plain, (spl68_256 | spl68_257 | spl68_258 | spl68_259 | spl68_260), inference(avatar_split_clause, [], [f284, f2026, f2022, f2018, f2014, f2010])).
fof(f284, plain, ((e24 = unit2) | (e23 = unit2) | (e22 = unit2) | (e21 = unit2) | (e20 = unit2)), inference(cnf_transformation, [], [f5])).
fof(f2008, plain, (spl68_251 | spl68_252 | spl68_253 | spl68_254 | spl68_255), inference(avatar_split_clause, [], [f249, f2005, f2001, f1997, f1993, f1989])).
fof(f249, plain, ((op2(e20, e20) = e24) | (op2(e20, e20) = e23) | (op2(e20, e20) = e22) | (op2(e20, e20) = e21) | (e20 = op2(e20, e20))), inference(cnf_transformation, [], [f4])).
fof(f4, plain, (((e24 = op2(e24, e24)) | (e23 = op2(e24, e24)) | (e22 = op2(e24, e24)) | (e21 = op2(e24, e24)) | (e20 = op2(e24, e24))) & ((e24 = op2(e24, e23)) | (e23 = op2(e24, e23)) | (e22 = op2(e24, e23)) | (e21 = op2(e24, e23)) | (e20 = op2(e24, e23))) & ((e24 = op2(e24, e22)) | (e23 = op2(e24, e22)) | (e22 = op2(e24, e22)) | (e21 = op2(e24, e22)) | (e20 = op2(e24, e22))) & ((e24 = op2(e24, e21)) | (e23 = op2(e24, e21)) | (e22 = op2(e24, e21)) | (e21 = op2(e24, e21)) | (e20 = op2(e24, e21))) & ((e24 = op2(e24, e20)) | (e23 = op2(e24, e20)) | (e22 = op2(e24, e20)) | (e21 = op2(e24, e20)) | (e20 = op2(e24, e20))) & ((e24 = op2(e23, e24)) | (e23 = op2(e23, e24)) | (e22 = op2(e23, e24)) | (e21 = op2(e23, e24)) | (e20 = op2(e23, e24))) & ((e24 = op2(e23, e23)) | (e23 = op2(e23, e23)) | (e22 = op2(e23, e23)) | (e21 = op2(e23, e23)) | (e20 = op2(e23, e23))) & ((e24 = op2(e23, e22)) | (e23 = op2(e23, e22)) | (e22 = op2(e23, e22)) | (e21 = op2(e23, e22)) | (e20 = op2(e23, e22))) & ((e24 = op2(e23, e21)) | (e23 = op2(e23, e21)) | (e22 = op2(e23, e21)) | (e21 = op2(e23, e21)) | (e20 = op2(e23, e21))) & ((e24 = op2(e23, e20)) | (e23 = op2(e23, e20)) | (e22 = op2(e23, e20)) | (e21 = op2(e23, e20)) | (e20 = op2(e23, e20))) & ((e24 = op2(e22, e24)) | (e23 = op2(e22, e24)) | (e22 = op2(e22, e24)) | (e21 = op2(e22, e24)) | (e20 = op2(e22, e24))) & ((e24 = op2(e22, e23)) | (e23 = op2(e22, e23)) | (e22 = op2(e22, e23)) | (e21 = op2(e22, e23)) | (e20 = op2(e22, e23))) & ((e24 = op2(e22, e22)) | (e23 = op2(e22, e22)) | (e22 = op2(e22, e22)) | (e21 = op2(e22, e22)) | (e20 = op2(e22, e22))) & ((e24 = op2(e22, e21)) | (e23 = op2(e22, e21)) | (e22 = op2(e22, e21)) | (e21 = op2(e22, e21)) | (e20 = op2(e22, e21))) & ((e24 = op2(e22, e20)) | (e23 = op2(e22, e20)) | (e22 = op2(e22, e20)) | (e21 = op2(e22, e20)) | (e20 = op2(e22, e20))) & ((e24 = op2(e21, e24)) | (e23 = op2(e21, e24)) | (e22 = op2(e21, e24)) | (e21 = op2(e21, e24)) | (e20 = op2(e21, e24))) & ((e24 = op2(e21, e23)) | (e23 = op2(e21, e23)) | (e22 = op2(e21, e23)) | (e21 = op2(e21, e23)) | (e20 = op2(e21, e23))) & ((e24 = op2(e21, e22)) | (e23 = op2(e21, e22)) | (e22 = op2(e21, e22)) | (e21 = op2(e21, e22)) | (e20 = op2(e21, e22))) & ((e24 = op2(e21, e21)) | (e23 = op2(e21, e21)) | (e22 = op2(e21, e21)) | (e21 = op2(e21, e21)) | (e20 = op2(e21, e21))) & ((e24 = op2(e21, e20)) | (e23 = op2(e21, e20)) | (e22 = op2(e21, e20)) | (e21 = op2(e21, e20)) | (e20 = op2(e21, e20))) & ((e24 = op2(e20, e24)) | (e23 = op2(e20, e24)) | (e22 = op2(e20, e24)) | (e21 = op2(e20, e24)) | (e20 = op2(e20, e24))) & ((e24 = op2(e20, e23)) | (e23 = op2(e20, e23)) | (e22 = op2(e20, e23)) | (e21 = op2(e20, e23)) | (e20 = op2(e20, e23))) & ((e24 = op2(e20, e22)) | (e23 = op2(e20, e22)) | (e22 = op2(e20, e22)) | (e21 = op2(e20, e22)) | (e20 = op2(e20, e22))) & ((e24 = op2(e20, e21)) | (e23 = op2(e20, e21)) | (e22 = op2(e20, e21)) | (e21 = op2(e20, e21)) | (e20 = op2(e20, e21))) & ((op2(e20, e20) = e24) | (op2(e20, e20) = e23) | (op2(e20, e20) = e22) | (op2(e20, e20) = e21) | (e20 = op2(e20, e20)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG052+1.p', ax4)).
fof(f1840, plain, (spl68_211 | spl68_212 | spl68_213 | spl68_214 | spl68_215), inference(avatar_split_clause, [], [f257, f1837, f1833, f1829, f1825, f1821])).
fof(f257, plain, ((e24 = op2(e21, e23)) | (e23 = op2(e21, e23)) | (e22 = op2(e21, e23)) | (e21 = op2(e21, e23)) | (e20 = op2(e21, e23))), inference(cnf_transformation, [], [f4])).
fof(f1714, plain, (spl68_181 | spl68_182 | spl68_183 | spl68_184 | spl68_185), inference(avatar_split_clause, [], [f263, f1711, f1707, f1703, f1699, f1695])).
fof(f263, plain, ((e24 = op2(e22, e24)) | (e23 = op2(e22, e24)) | (e22 = op2(e22, e24)) | (e21 = op2(e22, e24)) | (e20 = op2(e22, e24))), inference(cnf_transformation, [], [f4])).
fof(f1693, plain, (spl68_176 | spl68_177 | spl68_178 | spl68_179 | spl68_180), inference(avatar_split_clause, [], [f264, f1690, f1686, f1682, f1678, f1674])).
fof(f264, plain, ((e24 = op2(e23, e20)) | (e23 = op2(e23, e20)) | (e22 = op2(e23, e20)) | (e21 = op2(e23, e20)) | (e20 = op2(e23, e20))), inference(cnf_transformation, [], [f4])).
fof(f1672, plain, (spl68_171 | spl68_172 | spl68_173 | spl68_174 | spl68_175), inference(avatar_split_clause, [], [f265, f1669, f1665, f1661, f1657, f1653])).
fof(f265, plain, ((e24 = op2(e23, e21)) | (e23 = op2(e23, e21)) | (e22 = op2(e23, e21)) | (e21 = op2(e23, e21)) | (e20 = op2(e23, e21))), inference(cnf_transformation, [], [f4])).
fof(f1651, plain, (spl68_166 | spl68_167 | spl68_168 | spl68_169 | spl68_170), inference(avatar_split_clause, [], [f266, f1648, f1644, f1640, f1636, f1632])).
fof(f266, plain, ((e24 = op2(e23, e22)) | (e23 = op2(e23, e22)) | (e22 = op2(e23, e22)) | (e21 = op2(e23, e22)) | (e20 = op2(e23, e22))), inference(cnf_transformation, [], [f4])).
fof(f1630, plain, (spl68_161 | spl68_162 | spl68_163 | spl68_164 | spl68_165), inference(avatar_split_clause, [], [f267, f1627, f1623, f1619, f1615, f1611])).
fof(f267, plain, ((e24 = op2(e23, e23)) | (e23 = op2(e23, e23)) | (e22 = op2(e23, e23)) | (e21 = op2(e23, e23)) | (e20 = op2(e23, e23))), inference(cnf_transformation, [], [f4])).
fof(f1588, plain, (spl68_151 | spl68_152 | spl68_153 | spl68_154 | spl68_155), inference(avatar_split_clause, [], [f269, f1585, f1581, f1577, f1573, f1569])).
fof(f269, plain, ((e24 = op2(e24, e20)) | (e23 = op2(e24, e20)) | (e22 = op2(e24, e20)) | (e21 = op2(e24, e20)) | (e20 = op2(e24, e20))), inference(cnf_transformation, [], [f4])).
fof(f1546, plain, (spl68_141 | spl68_142 | spl68_143 | spl68_144 | spl68_145), inference(avatar_split_clause, [], [f271, f1543, f1539, f1535, f1531, f1527])).
fof(f271, plain, ((e24 = op2(e24, e22)) | (e23 = op2(e24, e22)) | (e22 = op2(e24, e22)) | (e21 = op2(e24, e22)) | (e20 = op2(e24, e22))), inference(cnf_transformation, [], [f4])).
fof(f1480, plain, (spl68_122 | spl68_97 | spl68_72 | spl68_47 | spl68_22), inference(avatar_split_clause, [], [f202, f977, f1082, f1187, f1292, f1397])).
fof(f202, plain, ((e11 = op1(e14, e10)) | (e11 = op1(e13, e10)) | (e11 = op1(e12, e10)) | (e11 = op1(e11, e10)) | (op1(e10, e10) = e11)), inference(cnf_transformation, [], [f3])).
fof(f3, plain, (((e14 = op1(e14, e14)) | (e14 = op1(e13, e14)) | (e14 = op1(e12, e14)) | (e14 = op1(e11, e14)) | (e14 = op1(e10, e14))) & ((e14 = op1(e14, e14)) | (e14 = op1(e14, e13)) | (e14 = op1(e14, e12)) | (e14 = op1(e14, e11)) | (e14 = op1(e14, e10))) & ((e13 = op1(e14, e14)) | (e13 = op1(e13, e14)) | (e13 = op1(e12, e14)) | (e13 = op1(e11, e14)) | (e13 = op1(e10, e14))) & ((e13 = op1(e14, e14)) | (e13 = op1(e14, e13)) | (e13 = op1(e14, e12)) | (e13 = op1(e14, e11)) | (e13 = op1(e14, e10))) & ((e12 = op1(e14, e14)) | (e12 = op1(e13, e14)) | (e12 = op1(e12, e14)) | (e12 = op1(e11, e14)) | (e12 = op1(e10, e14))) & ((e12 = op1(e14, e14)) | (e12 = op1(e14, e13)) | (e12 = op1(e14, e12)) | (e12 = op1(e14, e11)) | (e12 = op1(e14, e10))) & ((e11 = op1(e14, e14)) | (e11 = op1(e13, e14)) | (e11 = op1(e12, e14)) | (e11 = op1(e11, e14)) | (e11 = op1(e10, e14))) & ((e11 = op1(e14, e14)) | (e11 = op1(e14, e13)) | (e11 = op1(e14, e12)) | (e11 = op1(e14, e11)) | (e11 = op1(e14, e10))) & ((e10 = op1(e14, e14)) | (e10 = op1(e13, e14)) | (e10 = op1(e12, e14)) | (e10 = op1(e11, e14)) | (e10 = op1(e10, e14))) & ((e10 = op1(e14, e14)) | (e10 = op1(e14, e13)) | (e10 = op1(e14, e12)) | (e10 = op1(e14, e11)) | (e10 = op1(e14, e10))) & ((e14 = op1(e14, e13)) | (e14 = op1(e13, e13)) | (e14 = op1(e12, e13)) | (e14 = op1(e11, e13)) | (e14 = op1(e10, e13))) & ((e14 = op1(e13, e14)) | (e14 = op1(e13, e13)) | (e14 = op1(e13, e12)) | (e14 = op1(e13, e11)) | (e14 = op1(e13, e10))) & ((e13 = op1(e14, e13)) | (e13 = op1(e13, e13)) | (e13 = op1(e12, e13)) | (e13 = op1(e11, e13)) | (e13 = op1(e10, e13))) & ((e13 = op1(e13, e14)) | (e13 = op1(e13, e13)) | (e13 = op1(e13, e12)) | (e13 = op1(e13, e11)) | (e13 = op1(e13, e10))) & ((e12 = op1(e14, e13)) | (e12 = op1(e13, e13)) | (e12 = op1(e12, e13)) | (e12 = op1(e11, e13)) | (e12 = op1(e10, e13))) & ((e12 = op1(e13, e14)) | (e12 = op1(e13, e13)) | (e12 = op1(e13, e12)) | (e12 = op1(e13, e11)) | (e12 = op1(e13, e10))) & ((e11 = op1(e14, e13)) | (e11 = op1(e13, e13)) | (e11 = op1(e12, e13)) | (e11 = op1(e11, e13)) | (e11 = op1(e10, e13))) & ((e11 = op1(e13, e14)) | (e11 = op1(e13, e13)) | (e11 = op1(e13, e12)) | (e11 = op1(e13, e11)) | (e11 = op1(e13, e10))) & ((e10 = op1(e14, e13)) | (e10 = op1(e13, e13)) | (e10 = op1(e12, e13)) | (e10 = op1(e11, e13)) | (e10 = op1(e10, e13))) & ((e10 = op1(e13, e14)) | (e10 = op1(e13, e13)) | (e10 = op1(e13, e12)) | (e10 = op1(e13, e11)) | (e10 = op1(e13, e10))) & ((e14 = op1(e14, e12)) | (e14 = op1(e13, e12)) | (e14 = op1(e12, e12)) | (e14 = op1(e11, e12)) | (e14 = op1(e10, e12))) & ((e14 = op1(e12, e14)) | (e14 = op1(e12, e13)) | (e14 = op1(e12, e12)) | (e14 = op1(e12, e11)) | (e14 = op1(e12, e10))) & ((e13 = op1(e14, e12)) | (e13 = op1(e13, e12)) | (e13 = op1(e12, e12)) | (e13 = op1(e11, e12)) | (e13 = op1(e10, e12))) & ((e13 = op1(e12, e14)) | (e13 = op1(e12, e13)) | (e13 = op1(e12, e12)) | (e13 = op1(e12, e11)) | (e13 = op1(e12, e10))) & ((e12 = op1(e14, e12)) | (e12 = op1(e13, e12)) | (e12 = op1(e12, e12)) | (e12 = op1(e11, e12)) | (e12 = op1(e10, e12))) & ((e12 = op1(e12, e14)) | (e12 = op1(e12, e13)) | (e12 = op1(e12, e12)) | (e12 = op1(e12, e11)) | (e12 = op1(e12, e10))) & ((e11 = op1(e14, e12)) | (e11 = op1(e13, e12)) | (e11 = op1(e12, e12)) | (e11 = op1(e11, e12)) | (e11 = op1(e10, e12))) & ((e11 = op1(e12, e14)) | (e11 = op1(e12, e13)) | (e11 = op1(e12, e12)) | (e11 = op1(e12, e11)) | (e11 = op1(e12, e10))) & ((e10 = op1(e14, e12)) | (e10 = op1(e13, e12)) | (e10 = op1(e12, e12)) | (e10 = op1(e11, e12)) | (e10 = op1(e10, e12))) & ((e10 = op1(e12, e14)) | (e10 = op1(e12, e13)) | (e10 = op1(e12, e12)) | (e10 = op1(e12, e11)) | (e10 = op1(e12, e10))) & ((e14 = op1(e14, e11)) | (e14 = op1(e13, e11)) | (e14 = op1(e12, e11)) | (e14 = op1(e11, e11)) | (e14 = op1(e10, e11))) & ((e14 = op1(e11, e14)) | (e14 = op1(e11, e13)) | (e14 = op1(e11, e12)) | (e14 = op1(e11, e11)) | (e14 = op1(e11, e10))) & ((e13 = op1(e14, e11)) | (e13 = op1(e13, e11)) | (e13 = op1(e12, e11)) | (e13 = op1(e11, e11)) | (e13 = op1(e10, e11))) & ((e13 = op1(e11, e14)) | (e13 = op1(e11, e13)) | (e13 = op1(e11, e12)) | (e13 = op1(e11, e11)) | (e13 = op1(e11, e10))) & ((e12 = op1(e14, e11)) | (e12 = op1(e13, e11)) | (e12 = op1(e12, e11)) | (e12 = op1(e11, e11)) | (e12 = op1(e10, e11))) & ((e12 = op1(e11, e14)) | (e12 = op1(e11, e13)) | (e12 = op1(e11, e12)) | (e12 = op1(e11, e11)) | (e12 = op1(e11, e10))) & ((e11 = op1(e14, e11)) | (e11 = op1(e13, e11)) | (e11 = op1(e12, e11)) | (e11 = op1(e11, e11)) | (e11 = op1(e10, e11))) & ((e11 = op1(e11, e14)) | (e11 = op1(e11, e13)) | (e11 = op1(e11, e12)) | (e11 = op1(e11, e11)) | (e11 = op1(e11, e10))) & ((e10 = op1(e14, e11)) | (e10 = op1(e13, e11)) | (e10 = op1(e12, e11)) | (e10 = op1(e11, e11)) | (e10 = op1(e10, e11))) & ((e10 = op1(e11, e14)) | (e10 = op1(e11, e13)) | (e10 = op1(e11, e12)) | (e10 = op1(e11, e11)) | (e10 = op1(e11, e10))) & ((e14 = op1(e14, e10)) | (e14 = op1(e13, e10)) | (e14 = op1(e12, e10)) | (e14 = op1(e11, e10)) | (op1(e10, e10) = e14)) & ((e14 = op1(e10, e14)) | (e14 = op1(e10, e13)) | (e14 = op1(e10, e12)) | (e14 = op1(e10, e11)) | (op1(e10, e10) = e14)) & ((e13 = op1(e14, e10)) | (e13 = op1(e13, e10)) | (e13 = op1(e12, e10)) | (e13 = op1(e11, e10)) | (op1(e10, e10) = e13)) & ((e13 = op1(e10, e14)) | (e13 = op1(e10, e13)) | (e13 = op1(e10, e12)) | (e13 = op1(e10, e11)) | (op1(e10, e10) = e13)) & ((e12 = op1(e14, e10)) | (e12 = op1(e13, e10)) | (e12 = op1(e12, e10)) | (e12 = op1(e11, e10)) | (op1(e10, e10) = e12)) & ((e12 = op1(e10, e14)) | (e12 = op1(e10, e13)) | (e12 = op1(e10, e12)) | (e12 = op1(e10, e11)) | (op1(e10, e10) = e12)) & ((e11 = op1(e14, e10)) | (e11 = op1(e13, e10)) | (e11 = op1(e12, e10)) | (e11 = op1(e11, e10)) | (op1(e10, e10) = e11)) & ((e11 = op1(e10, e14)) | (e11 = op1(e10, e13)) | (e11 = op1(e10, e12)) | (e11 = op1(e10, e11)) | (op1(e10, e10) = e11)) & ((e10 = op1(e14, e10)) | (e10 = op1(e13, e10)) | (e10 = op1(e12, e10)) | (e10 = op1(e11, e10)) | (e10 = op1(e10, e10))) & ((e10 = op1(e10, e14)) | (e10 = op1(e10, e13)) | (e10 = op1(e10, e12)) | (e10 = op1(e10, e11)) | (e10 = op1(e10, e10)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG052+1.p', ax3)).
fof(f1475, plain, (spl68_125 | spl68_120 | spl68_115 | spl68_110 | spl68_105), inference(avatar_split_clause, [], [f207, f1325, f1346, f1367, f1388, f1409])).
fof(f207, plain, ((e14 = op1(e10, e14)) | (e14 = op1(e10, e13)) | (e14 = op1(e10, e12)) | (e14 = op1(e10, e11)) | (op1(e10, e10) = e14)), inference(cnf_transformation, [], [f3])).
fof(f1472, plain, (spl68_116 | spl68_91 | spl68_66 | spl68_41 | spl68_16), inference(avatar_split_clause, [], [f210, f952, f1057, f1162, f1267, f1372])).
fof(f210, plain, ((e10 = op1(e14, e11)) | (e10 = op1(e13, e11)) | (e10 = op1(e12, e11)) | (e10 = op1(e11, e11)) | (e10 = op1(e10, e11))), inference(cnf_transformation, [], [f3])).
fof(f1468, plain, (spl68_118 | spl68_93 | spl68_68 | spl68_43 | spl68_18), inference(avatar_split_clause, [], [f214, f960, f1065, f1170, f1275, f1380])).
fof(f214, plain, ((e12 = op1(e14, e11)) | (e12 = op1(e13, e11)) | (e12 = op1(e12, e11)) | (e12 = op1(e11, e11)) | (e12 = op1(e10, e11))), inference(cnf_transformation, [], [f3])).
fof(f1467, plain, (spl68_99 | spl68_94 | spl68_89 | spl68_84 | spl68_79), inference(avatar_split_clause, [], [f215, f1216, f1237, f1258, f1279, f1300])).
fof(f215, plain, ((e13 = op1(e11, e14)) | (e13 = op1(e11, e13)) | (e13 = op1(e11, e12)) | (e13 = op1(e11, e11)) | (e13 = op1(e11, e10))), inference(cnf_transformation, [], [f3])).
fof(f1452, plain, (spl68_106 | spl68_81 | spl68_56 | spl68_31 | spl68_6), inference(avatar_split_clause, [], [f230, f910, f1015, f1120, f1225, f1330])).
fof(f230, plain, ((e10 = op1(e14, e13)) | (e10 = op1(e13, e13)) | (e10 = op1(e12, e13)) | (e10 = op1(e11, e13)) | (e10 = op1(e10, e13))), inference(cnf_transformation, [], [f3])).
fof(f1451, plain, (spl68_47 | spl68_42 | spl68_37 | spl68_32 | spl68_27), inference(avatar_split_clause, [], [f231, f998, f1019, f1040, f1061, f1082])).
fof(f231, plain, ((e11 = op1(e13, e14)) | (e11 = op1(e13, e13)) | (e11 = op1(e13, e12)) | (e11 = op1(e13, e11)) | (e11 = op1(e13, e10))), inference(cnf_transformation, [], [f3])).
fof(f1450, plain, (spl68_107 | spl68_82 | spl68_57 | spl68_32 | spl68_7), inference(avatar_split_clause, [], [f232, f914, f1019, f1124, f1229, f1334])).
fof(f232, plain, ((e11 = op1(e14, e13)) | (e11 = op1(e13, e13)) | (e11 = op1(e12, e13)) | (e11 = op1(e11, e13)) | (e11 = op1(e10, e13))), inference(cnf_transformation, [], [f3])).
fof(f1445, plain, (spl68_50 | spl68_45 | spl68_40 | spl68_35 | spl68_30), inference(avatar_split_clause, [], [f237, f1010, f1031, f1052, f1073, f1094])).
fof(f237, plain, ((e14 = op1(e13, e14)) | (e14 = op1(e13, e13)) | (e14 = op1(e13, e12)) | (e14 = op1(e13, e11)) | (e14 = op1(e13, e10))), inference(cnf_transformation, [], [f3])).
fof(f1439, plain, (spl68_23 | spl68_18 | spl68_13 | spl68_8 | spl68_3), inference(avatar_split_clause, [], [f243, f897, f918, f939, f960, f981])).
fof(f243, plain, ((e12 = op1(e14, e14)) | (e12 = op1(e14, e13)) | (e12 = op1(e14, e12)) | (e12 = op1(e14, e11)) | (e12 = op1(e14, e10))), inference(cnf_transformation, [], [f3])).
fof(f1436, plain, (spl68_104 | spl68_79 | spl68_54 | spl68_29 | spl68_4), inference(avatar_split_clause, [], [f246, f901, f1006, f1111, f1216, f1321])).
fof(f246, plain, ((e13 = op1(e14, e14)) | (e13 = op1(e13, e14)) | (e13 = op1(e12, e14)) | (e13 = op1(e11, e14)) | (e13 = op1(e10, e14))), inference(cnf_transformation, [], [f3])).
fof(f1433, plain, (spl68_126 | spl68_127 | spl68_128 | spl68_129 | spl68_130), inference(avatar_split_clause, [], [f198, f1430, f1426, f1422, f1418, f1414])).
fof(f198, plain, ((e14 = unit1) | (e13 = unit1) | (e12 = unit1) | (e11 = unit1) | (e10 = unit1)), inference(cnf_transformation, [], [f2])).
fof(f1412, plain, (spl68_121 | spl68_122 | spl68_123 | spl68_124 | spl68_125), inference(avatar_split_clause, [], [f163, f1409, f1405, f1401, f1397, f1393])).
fof(f163, plain, ((op1(e10, e10) = e14) | (op1(e10, e10) = e13) | (op1(e10, e10) = e12) | (op1(e10, e10) = e11) | (e10 = op1(e10, e10))), inference(cnf_transformation, [], [f1])).
fof(f1, plain, (((e14 = op1(e14, e14)) | (e13 = op1(e14, e14)) | (e12 = op1(e14, e14)) | (e11 = op1(e14, e14)) | (e10 = op1(e14, e14))) & ((e14 = op1(e14, e13)) | (e13 = op1(e14, e13)) | (e12 = op1(e14, e13)) | (e11 = op1(e14, e13)) | (e10 = op1(e14, e13))) & ((e14 = op1(e14, e12)) | (e13 = op1(e14, e12)) | (e12 = op1(e14, e12)) | (e11 = op1(e14, e12)) | (e10 = op1(e14, e12))) & ((e14 = op1(e14, e11)) | (e13 = op1(e14, e11)) | (e12 = op1(e14, e11)) | (e11 = op1(e14, e11)) | (e10 = op1(e14, e11))) & ((e14 = op1(e14, e10)) | (e13 = op1(e14, e10)) | (e12 = op1(e14, e10)) | (e11 = op1(e14, e10)) | (e10 = op1(e14, e10))) & ((e14 = op1(e13, e14)) | (e13 = op1(e13, e14)) | (e12 = op1(e13, e14)) | (e11 = op1(e13, e14)) | (e10 = op1(e13, e14))) & ((e14 = op1(e13, e13)) | (e13 = op1(e13, e13)) | (e12 = op1(e13, e13)) | (e11 = op1(e13, e13)) | (e10 = op1(e13, e13))) & ((e14 = op1(e13, e12)) | (e13 = op1(e13, e12)) | (e12 = op1(e13, e12)) | (e11 = op1(e13, e12)) | (e10 = op1(e13, e12))) & ((e14 = op1(e13, e11)) | (e13 = op1(e13, e11)) | (e12 = op1(e13, e11)) | (e11 = op1(e13, e11)) | (e10 = op1(e13, e11))) & ((e14 = op1(e13, e10)) | (e13 = op1(e13, e10)) | (e12 = op1(e13, e10)) | (e11 = op1(e13, e10)) | (e10 = op1(e13, e10))) & ((e14 = op1(e12, e14)) | (e13 = op1(e12, e14)) | (e12 = op1(e12, e14)) | (e11 = op1(e12, e14)) | (e10 = op1(e12, e14))) & ((e14 = op1(e12, e13)) | (e13 = op1(e12, e13)) | (e12 = op1(e12, e13)) | (e11 = op1(e12, e13)) | (e10 = op1(e12, e13))) & ((e14 = op1(e12, e12)) | (e13 = op1(e12, e12)) | (e12 = op1(e12, e12)) | (e11 = op1(e12, e12)) | (e10 = op1(e12, e12))) & ((e14 = op1(e12, e11)) | (e13 = op1(e12, e11)) | (e12 = op1(e12, e11)) | (e11 = op1(e12, e11)) | (e10 = op1(e12, e11))) & ((e14 = op1(e12, e10)) | (e13 = op1(e12, e10)) | (e12 = op1(e12, e10)) | (e11 = op1(e12, e10)) | (e10 = op1(e12, e10))) & ((e14 = op1(e11, e14)) | (e13 = op1(e11, e14)) | (e12 = op1(e11, e14)) | (e11 = op1(e11, e14)) | (e10 = op1(e11, e14))) & ((e14 = op1(e11, e13)) | (e13 = op1(e11, e13)) | (e12 = op1(e11, e13)) | (e11 = op1(e11, e13)) | (e10 = op1(e11, e13))) & ((e14 = op1(e11, e12)) | (e13 = op1(e11, e12)) | (e12 = op1(e11, e12)) | (e11 = op1(e11, e12)) | (e10 = op1(e11, e12))) & ((e14 = op1(e11, e11)) | (e13 = op1(e11, e11)) | (e12 = op1(e11, e11)) | (e11 = op1(e11, e11)) | (e10 = op1(e11, e11))) & ((e14 = op1(e11, e10)) | (e13 = op1(e11, e10)) | (e12 = op1(e11, e10)) | (e11 = op1(e11, e10)) | (e10 = op1(e11, e10))) & ((e14 = op1(e10, e14)) | (e13 = op1(e10, e14)) | (e12 = op1(e10, e14)) | (e11 = op1(e10, e14)) | (e10 = op1(e10, e14))) & ((e14 = op1(e10, e13)) | (e13 = op1(e10, e13)) | (e12 = op1(e10, e13)) | (e11 = op1(e10, e13)) | (e10 = op1(e10, e13))) & ((e14 = op1(e10, e12)) | (e13 = op1(e10, e12)) | (e12 = op1(e10, e12)) | (e11 = op1(e10, e12)) | (e10 = op1(e10, e12))) & ((e14 = op1(e10, e11)) | (e13 = op1(e10, e11)) | (e12 = op1(e10, e11)) | (e11 = op1(e10, e11)) | (e10 = op1(e10, e11))) & ((op1(e10, e10) = e14) | (op1(e10, e10) = e13) | (op1(e10, e10) = e12) | (op1(e10, e10) = e11) | (e10 = op1(e10, e10)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG052+1.p', ax1)).
fof(f1391, plain, (spl68_116 | spl68_117 | spl68_118 | spl68_119 | spl68_120), inference(avatar_split_clause, [], [f164, f1388, f1384, f1380, f1376, f1372])).
fof(f164, plain, ((e14 = op1(e10, e11)) | (e13 = op1(e10, e11)) | (e12 = op1(e10, e11)) | (e11 = op1(e10, e11)) | (e10 = op1(e10, e11))), inference(cnf_transformation, [], [f1])).
fof(f1244, plain, (spl68_81 | spl68_82 | spl68_83 | spl68_84 | spl68_85), inference(avatar_split_clause, [], [f171, f1241, f1237, f1233, f1229, f1225])).
fof(f171, plain, ((e14 = op1(e11, e13)) | (e13 = op1(e11, e13)) | (e12 = op1(e11, e13)) | (e11 = op1(e11, e13)) | (e10 = op1(e11, e13))), inference(cnf_transformation, [], [f1])).
fof(f1181, plain, (spl68_66 | spl68_67 | spl68_68 | spl68_69 | spl68_70), inference(avatar_split_clause, [], [f174, f1178, f1174, f1170, f1166, f1162])).
fof(f174, plain, ((e14 = op1(e12, e11)) | (e13 = op1(e12, e11)) | (e12 = op1(e12, e11)) | (e11 = op1(e12, e11)) | (e10 = op1(e12, e11))), inference(cnf_transformation, [], [f1])).
fof(f1118, plain, (spl68_51 | spl68_52 | spl68_53 | spl68_54 | spl68_55), inference(avatar_split_clause, [], [f177, f1115, f1111, f1107, f1103, f1099])).
fof(f177, plain, ((e14 = op1(e12, e14)) | (e13 = op1(e12, e14)) | (e12 = op1(e12, e14)) | (e11 = op1(e12, e14)) | (e10 = op1(e12, e14))), inference(cnf_transformation, [], [f1])).
fof(f1097, plain, (spl68_46 | spl68_47 | spl68_48 | spl68_49 | spl68_50), inference(avatar_split_clause, [], [f178, f1094, f1090, f1086, f1082, f1078])).
fof(f178, plain, ((e14 = op1(e13, e10)) | (e13 = op1(e13, e10)) | (e12 = op1(e13, e10)) | (e11 = op1(e13, e10)) | (e10 = op1(e13, e10))), inference(cnf_transformation, [], [f1])).
fof(f1055, plain, (spl68_36 | spl68_37 | spl68_38 | spl68_39 | spl68_40), inference(avatar_split_clause, [], [f180, f1052, f1048, f1044, f1040, f1036])).
fof(f180, plain, ((e14 = op1(e13, e12)) | (e13 = op1(e13, e12)) | (e12 = op1(e13, e12)) | (e11 = op1(e13, e12)) | (e10 = op1(e13, e12))), inference(cnf_transformation, [], [f1])).
fof(f1034, plain, (spl68_31 | spl68_32 | spl68_33 | spl68_34 | spl68_35), inference(avatar_split_clause, [], [f181, f1031, f1027, f1023, f1019, f1015])).
fof(f181, plain, ((e14 = op1(e13, e13)) | (e13 = op1(e13, e13)) | (e12 = op1(e13, e13)) | (e11 = op1(e13, e13)) | (e10 = op1(e13, e13))), inference(cnf_transformation, [], [f1])).
fof(f950, plain, (spl68_11 | spl68_12 | spl68_13 | spl68_14 | spl68_15), inference(avatar_split_clause, [], [f185, f947, f943, f939, f935, f931])).
fof(f185, plain, ((e14 = op1(e14, e12)) | (e13 = op1(e14, e12)) | (e12 = op1(e14, e12)) | (e11 = op1(e14, e12)) | (e10 = op1(e14, e12))), inference(cnf_transformation, [], [f1])).