fof(f43645, plain, $false, inference(avatar_sat_refutation, [], [f2820, f2920, f3270, f3320, f3395, f3420, f3670, f3795, f3995, f4070, f4095, f4445, f4470, f4745, f4770, f4820, f4845, f4870, f4927, f4939, f4941, f4943, f4944, f5078, f5080, f5082, f5084, f7988, f8003, f8019, f8039, f15272, f18033, f18034, f18036, f18297, f18299, f18301, f18302, f18303, f18304, f18318, f18319, f18322, f18335, f18385, f18396, f18510, f18520, f18532, f18534, f18551, f18571, f18572, f18575, f18611, f18681, f18968, f18969, f18970, f18971, f18972, f18973, f18974, f18975, f18976, f18977, f18989, f19739, f19830, f19901, f19904, f19918, f19945, f19951, f19989, f19991, f19997, f20106, f20107, f20130, f20131, f20133, f20146, f20150, f20153, f20164, f20173, f20363, f20614, f20615, f20617, f20636, f20642, f20661, f20663, f20669, f20685, f20690, f20777, f21030, f21705, f21745, f21775, f22333, f22340, f22392, f22710, f22781, f22833, f22842, f22850, f22885, f22897, f22908, f22933, f24134, f24209, f24267, f24568, f24730, f25224, f25489, f25817, f25823, f26019, f26102, f26646, f26684, f26732, f27125, f27437, f27667, f27704, f27817, f27911, f28527, f28968, f29066, f29092, f29262, f29410, f29453, f29529, f29635, f29657, f29668, f29781, f30042, f30054, f30099, f30450, f30487, f30566, f31026, f31704, f32097, f32366, f32411, f32545, f32585, f32822, f32825, f33059, f33060, f33091, f33619, f33632, f33961, f34086, f34091, f34096, f34121, f34488, f34503, f34619, f34787, f35096, f35108, f36097, f36505, f36584, f36634, f37039, f37230, f37309, f37742, f38165, f38402, f38403, f38404, f38616, f38646, f38728, f38757, f38758, f39053, f39074, f39202, f39770, f40044, f40273, f40398, f40411, f40427, f40473, f40520, f40561, f40628, f40653, f40660, f40743, f40865, f40868, f41271, f41387, f41403, f41404, f41716, f41817, f41926, f42051, f42173, f42285, f42379, f42441, f42502, f42562, f42617, f42670, f42723, f42773, f42826, f42873, f42922, f42964, f43003, f43049, f43096, f43138, f43178, f43215, f43257, f43301, f43343, f43392, f43436, f43479, f43525, f43567, f43609])).
fof(f43609, plain, (~ spl150_206 | ~ spl150_1039 | spl150_1972), inference(avatar_contradiction_clause, [], [f43608])).
fof(f43608, plain, ($false | (~ spl150_206 | ~ spl150_1039 | spl150_1972)), inference(subsumption_resolution, [], [f43607, f8448])).
fof(f8448, plain, ((e21 = h6(e13)) | ~ spl150_1039), inference(avatar_component_clause, [], [f8447])).
fof(f8447, plain, (spl150_1039 <=> (e21 = h6(e13))), introduced(avatar_definition, [new_symbols(naming, [spl150_1039])])).
fof(f43607, plain, (~ (e21 = h6(e13)) | (~ spl150_206 | spl150_1972)), inference(forward_demodulation, [], [f43606, f1536])).
fof(f1536, plain, (e21 = h10(e11)), inference(cnf_transformation, [], [f35])).
fof(f35, plain, ((op2(op2(e25, e21), op2(e25, e21)) = h10(e14)) & (op2(e25, e21) = h10(e13)) & (op2(e25, op2(op2(e25, e21), op2(e25, e21))) = h10(e12)) & (op2(op2(e25, op2(op2(e25, e21), op2(e25, e21))), op2(e25, op2(op2(e25, e21), op2(e25, e21)))) = h10(e10)) & (e25 = h10(e15)) & (e21 = h10(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG027+1.p', ax35)).
fof(f43606, plain, (~ (h6(e13) = h10(e11)) | (~ spl150_206 | spl150_1972)), inference(forward_demodulation, [], [f15267, f3603])).
fof(f3603, plain, ((e11 = op1(e10, e11)) | ~ spl150_206), inference(avatar_component_clause, [], [f3601])).
fof(f3601, plain, (spl150_206 <=> (e11 = op1(e10, e11))), introduced(avatar_definition, [new_symbols(naming, [spl150_206])])).
fof(f15267, plain, (~ (h6(e13) = h10(op1(e10, e11))) | spl150_1972), inference(avatar_component_clause, [], [f15265])).
fof(f15265, plain, (spl150_1972 <=> (h6(e13) = h10(op1(e10, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl150_1972])])).
fof(f43567, plain, (~ spl150_191 | ~ spl150_725 | spl150_1969), inference(avatar_contradiction_clause, [], [f43566])).
fof(f43566, plain, ($false | (~ spl150_191 | ~ spl150_725 | spl150_1969)), inference(subsumption_resolution, [], [f43565, f5319])).
fof(f5319, plain, (e24 = h10(e14)), inference(forward_demodulation, [], [f1541, f1481])).
fof(f1481, plain, (e24 = op2(op2(e25, e21), op2(e25, e21))), inference(cnf_transformation, [], [f25])).
fof(f25, plain, ((e24 = op2(op2(e25, e21), op2(e25, e21))) & (e23 = op2(e25, e21)) & (e22 = op2(e25, op2(op2(e25, e21), op2(e25, e21)))) & (e20 = op2(op2(e25, op2(op2(e25, e21), op2(e25, e21))), op2(e25, op2(op2(e25, e21), op2(e25, e21)))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG027+1.p', ax25)).
fof(f1541, plain, (op2(op2(e25, e21), op2(e25, e21)) = h10(e14)), inference(cnf_transformation, [], [f35])).
fof(f43565, plain, (~ (e24 = h10(e14)) | (~ spl150_191 | ~ spl150_725 | spl150_1969)), inference(forward_demodulation, [], [f43564, f6772])).
fof(f6772, plain, ((e24 = h21(e13)) | ~ spl150_725), inference(avatar_component_clause, [], [f6771])).
fof(f6771, plain, (spl150_725 <=> (e24 = h21(e13))), introduced(avatar_definition, [new_symbols(naming, [spl150_725])])).
fof(f43564, plain, (~ (h10(e14) = h21(e13)) | (~ spl150_191 | spl150_1969)), inference(forward_demodulation, [], [f15255, f3540])).
fof(f3540, plain, ((e14 = op1(e10, e14)) | ~ spl150_191), inference(avatar_component_clause, [], [f3538])).
fof(f3538, plain, (spl150_191 <=> (e14 = op1(e10, e14))), introduced(avatar_definition, [new_symbols(naming, [spl150_191])])).
fof(f15255, plain, (~ (h21(e13) = h10(op1(e10, e14))) | spl150_1969), inference(avatar_component_clause, [], [f15253])).
fof(f15253, plain, (spl150_1969 <=> (h21(e13) = h10(op1(e10, e14)))), introduced(avatar_definition, [new_symbols(naming, [spl150_1969])])).
fof(f43525, plain, (~ spl150_186 | ~ spl150_1336 | spl150_1968), inference(avatar_contradiction_clause, [], [f43524])).
fof(f43524, plain, ($false | (~ spl150_186 | ~ spl150_1336 | spl150_1968)), inference(subsumption_resolution, [], [f43523, f1537])).
fof(f1537, plain, (e25 = h10(e15)), inference(cnf_transformation, [], [f35])).
fof(f43523, plain, (~ (e25 = h10(e15)) | (~ spl150_186 | ~ spl150_1336 | spl150_1968)), inference(forward_demodulation, [], [f43522, f10450])).
fof(f10450, plain, ((e25 = h26(e13)) | ~ spl150_1336), inference(avatar_component_clause, [], [f10449])).
fof(f10449, plain, (spl150_1336 <=> (e25 = h26(e13))), introduced(avatar_definition, [new_symbols(naming, [spl150_1336])])).
fof(f43522, plain, (~ (h10(e15) = h26(e13)) | (~ spl150_186 | spl150_1968)), inference(forward_demodulation, [], [f15251, f3519])).
fof(f3519, plain, ((e15 = op1(e10, e15)) | ~ spl150_186), inference(avatar_component_clause, [], [f3517])).
fof(f3517, plain, (spl150_186 <=> (e15 = op1(e10, e15))), introduced(avatar_definition, [new_symbols(naming, [spl150_186])])).
fof(f15251, plain, (~ (h26(e13) = h10(op1(e10, e15))) | spl150_1968), inference(avatar_component_clause, [], [f15249])).
fof(f15249, plain, (spl150_1968 <=> (h26(e13) = h10(op1(e10, e15)))), introduced(avatar_definition, [new_symbols(naming, [spl150_1968])])).
fof(f43479, plain, (~ spl150_176 | ~ spl150_434 | spl150_1967), inference(avatar_contradiction_clause, [], [f43478])).
fof(f43478, plain, ($false | (~ spl150_176 | ~ spl150_434 | spl150_1967)), inference(subsumption_resolution, [], [f43477, f39771])).
fof(f39771, plain, ((e21 = h1(e13)) | ~ spl150_434), inference(forward_demodulation, [], [f1486, f4553])).
fof(f4553, plain, ((e21 = op2(e21, e20)) | ~ spl150_434), inference(avatar_component_clause, [], [f4551])).
fof(f4551, plain, (spl150_434 <=> (e21 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl150_434])])).
fof(f1486, plain, (op2(e21, e20) = h1(e13)), inference(cnf_transformation, [], [f26])).
fof(f26, plain, ((op2(op2(e21, e20), op2(e21, e20)) = h1(e14)) & (op2(e21, e20) = h1(e13)) & (op2(e21, op2(op2(e21, e20), op2(e21, e20))) = h1(e12)) & (h1(e10) = op2(op2(e21, op2(op2(e21, e20), op2(e21, e20))), op2(e21, op2(op2(e21, e20), op2(e21, e20))))) & (e21 = h1(e15)) & (e20 = h1(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG027+1.p', ax26)).
fof(f43477, plain, (~ (e21 = h1(e13)) | (~ spl150_176 | spl150_1967)), inference(forward_demodulation, [], [f43476, f1536])).
fof(f43476, plain, (~ (h1(e13) = h10(e11)) | (~ spl150_176 | spl150_1967)), inference(forward_demodulation, [], [f15247, f3478])).
fof(f3478, plain, ((e11 = op1(e11, e10)) | ~ spl150_176), inference(avatar_component_clause, [], [f3476])).
fof(f3476, plain, (spl150_176 <=> (e11 = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl150_176])])).
fof(f15247, plain, (~ (h1(e13) = h10(op1(e11, e10))) | spl150_1967), inference(avatar_component_clause, [], [f15245])).
fof(f15245, plain, (spl150_1967 <=> (h1(e13) = h10(op1(e11, e10)))), introduced(avatar_definition, [new_symbols(naming, [spl150_1967])])).
fof(f43436, plain, (~ spl150_169 | ~ spl150_427 | spl150_1966), inference(avatar_contradiction_clause, [], [f43435])).
fof(f43435, plain, ($false | (~ spl150_169 | ~ spl150_427 | spl150_1966)), inference(subsumption_resolution, [], [f43434, f4524])).
fof(f4524, plain, ((e20 = op2(e21, e21)) | ~ spl150_427), inference(avatar_component_clause, [], [f4522])).
fof(f4522, plain, (spl150_427 <=> (e20 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl150_427])])).
fof(f43434, plain, (~ (e20 = op2(e21, e21)) | (~ spl150_169 | spl150_1966)), inference(forward_demodulation, [], [f43433, f5365])).
fof(f5365, plain, (e20 = h10(e10)), inference(forward_demodulation, [], [f5364, f5358])).
fof(f5358, plain, (e20 = op2(h10(e12), h10(e12))), inference(backward_demodulation, [], [f5085, f5339])).
fof(f5339, plain, (op2(e25, e24) = h10(e12)), inference(forward_demodulation, [], [f5338, f5336])).
fof(f5336, plain, (e24 = op2(h10(e13), h10(e13))), inference(backward_demodulation, [], [f1481, f1540])).
fof(f1540, plain, (op2(e25, e21) = h10(e13)), inference(cnf_transformation, [], [f35])).
fof(f5338, plain, (h10(e12) = op2(e25, op2(h10(e13), h10(e13)))), inference(forward_demodulation, [], [f1539, f1540])).
fof(f1539, plain, (op2(e25, op2(op2(e25, e21), op2(e25, e21))) = h10(e12)), inference(cnf_transformation, [], [f35])).
fof(f5085, plain, (e20 = op2(op2(e25, e24), op2(e25, e24))), inference(forward_demodulation, [], [f1478, f1481])).
fof(f1478, plain, (e20 = op2(op2(e25, op2(op2(e25, e21), op2(e25, e21))), op2(e25, op2(op2(e25, e21), op2(e25, e21))))), inference(cnf_transformation, [], [f25])).
fof(f5364, plain, (h10(e10) = op2(h10(e12), h10(e12))), inference(forward_demodulation, [], [f5363, f5339])).
fof(f5363, plain, (h10(e10) = op2(op2(e25, e24), op2(e25, e24))), inference(forward_demodulation, [], [f5362, f5336])).
fof(f5362, plain, (h10(e10) = op2(op2(e25, op2(h10(e13), h10(e13))), op2(e25, op2(h10(e13), h10(e13))))), inference(forward_demodulation, [], [f1538, f1540])).
fof(f1538, plain, (op2(op2(e25, op2(op2(e25, e21), op2(e25, e21))), op2(e25, op2(op2(e25, e21), op2(e25, e21)))) = h10(e10)), inference(cnf_transformation, [], [f35])).
fof(f43433, plain, (~ (op2(e21, e21) = h10(e10)) | (~ spl150_169 | spl150_1966)), inference(forward_demodulation, [], [f15243, f3449])).
fof(f3449, plain, ((e10 = op1(e11, e11)) | ~ spl150_169), inference(avatar_component_clause, [], [f3447])).
fof(f3447, plain, (spl150_169 <=> (e10 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl150_169])])).
fof(f15243, plain, (~ (op2(e21, e21) = h10(op1(e11, e11))) | spl150_1966), inference(avatar_component_clause, [], [f15241])).
fof(f15241, plain, (spl150_1966 <=> (op2(e21, e21) = h10(op1(e11, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl150_1966])])).
fof(f43392, plain, (~ spl150_149 | ~ spl150_609 | spl150_1962), inference(avatar_contradiction_clause, [], [f43391])).
fof(f43391, plain, ($false | (~ spl150_149 | ~ spl150_609 | spl150_1962)), inference(subsumption_resolution, [], [f43390, f5319])).
fof(f43390, plain, (~ (e24 = h10(e14)) | (~ spl150_149 | ~ spl150_609 | spl150_1962)), inference(forward_demodulation, [], [f43389, f6141])).
fof(f6141, plain, ((e24 = h27(e13)) | ~ spl150_609), inference(avatar_component_clause, [], [f6140])).
fof(f6140, plain, (spl150_609 <=> (e24 = h27(e13))), introduced(avatar_definition, [new_symbols(naming, [spl150_609])])).
fof(f43389, plain, (~ (h10(e14) = h27(e13)) | (~ spl150_149 | spl150_1962)), inference(forward_demodulation, [], [f15227, f3365])).
fof(f3365, plain, ((e14 = op1(e11, e15)) | ~ spl150_149), inference(avatar_component_clause, [], [f3363])).
fof(f3363, plain, (spl150_149 <=> (e14 = op1(e11, e15))), introduced(avatar_definition, [new_symbols(naming, [spl150_149])])).
fof(f15227, plain, (~ (h27(e13) = h10(op1(e11, e15))) | spl150_1962), inference(avatar_component_clause, [], [f15225])).
fof(f15225, plain, (spl150_1962 <=> (h27(e13) = h10(op1(e11, e15)))), introduced(avatar_definition, [new_symbols(naming, [spl150_1962])])).
fof(f43343, plain, (~ spl150_201 | ~ spl150_459 | ~ spl150_656 | spl150_1971), inference(avatar_contradiction_clause, [], [f43342])).
fof(f43342, plain, ($false | (~ spl150_201 | ~ spl150_459 | ~ spl150_656 | spl150_1971)), inference(subsumption_resolution, [], [f43341, f4657])).
fof(f4657, plain, ((e22 = op2(e20, e22)) | ~ spl150_459), inference(avatar_component_clause, [], [f4655])).
fof(f4655, plain, (spl150_459 <=> (e22 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl150_459])])).
fof(f43341, plain, (~ (e22 = op2(e20, e22)) | (~ spl150_201 | ~ spl150_656 | spl150_1971)), inference(forward_demodulation, [], [f43340, f6399])).
fof(f6399, plain, ((e22 = h10(e12)) | ~ spl150_656), inference(avatar_component_clause, [], [f6398])).
fof(f6398, plain, (spl150_656 <=> (e22 = h10(e12))), introduced(avatar_definition, [new_symbols(naming, [spl150_656])])).
fof(f43340, plain, (~ (op2(e20, e22) = h10(e12)) | (~ spl150_201 | ~ spl150_656 | spl150_1971)), inference(forward_demodulation, [], [f43339, f3582])).
fof(f3582, plain, ((e12 = op1(e10, e12)) | ~ spl150_201), inference(avatar_component_clause, [], [f3580])).
fof(f3580, plain, (spl150_201 <=> (e12 = op1(e10, e12))), introduced(avatar_definition, [new_symbols(naming, [spl150_201])])).
fof(f43339, plain, (~ (op2(e20, e22) = h10(op1(e10, e12))) | (~ spl150_656 | spl150_1971)), inference(forward_demodulation, [], [f15263, f6399])).
fof(f15263, plain, (~ (h10(op1(e10, e12)) = op2(e20, h10(e12))) | spl150_1971), inference(avatar_component_clause, [], [f15261])).
fof(f15261, plain, (spl150_1971 <=> (h10(op1(e10, e12)) = op2(e20, h10(e12)))), introduced(avatar_definition, [new_symbols(naming, [spl150_1971])])).
fof(f43301, plain, (~ spl150_156 | ~ spl150_1495 | spl150_1963), inference(avatar_contradiction_clause, [], [f43300])).
fof(f43300, plain, ($false | (~ spl150_156 | ~ spl150_1495 | spl150_1963)), inference(subsumption_resolution, [], [f43299, f1537])).
fof(f43299, plain, (~ (e25 = h10(e15)) | (~ spl150_156 | ~ spl150_1495 | spl150_1963)), inference(forward_demodulation, [], [f43298, f11819])).
fof(f11819, plain, ((e25 = h22(e13)) | ~ spl150_1495), inference(avatar_component_clause, [], [f11818])).
fof(f11818, plain, (spl150_1495 <=> (e25 = h22(e13))), introduced(avatar_definition, [new_symbols(naming, [spl150_1495])])).
fof(f43298, plain, (~ (h10(e15) = h22(e13)) | (~ spl150_156 | spl150_1963)), inference(forward_demodulation, [], [f15231, f3394])).
fof(f3394, plain, ((e15 = op1(e11, e14)) | ~ spl150_156), inference(avatar_component_clause, [], [f3392])).
fof(f3392, plain, (spl150_156 <=> (e15 = op1(e11, e14))), introduced(avatar_definition, [new_symbols(naming, [spl150_156])])).
fof(f15231, plain, (~ (h22(e13) = h10(op1(e11, e14))) | spl150_1963), inference(avatar_component_clause, [], [f15229])).
fof(f15229, plain, (spl150_1963 <=> (h22(e13) = h10(op1(e11, e14)))), introduced(avatar_definition, [new_symbols(naming, [spl150_1963])])).
fof(f43257, plain, (~ spl150_141 | ~ spl150_267 | ~ spl150_329 | ~ spl150_656 | spl150_1961), inference(avatar_contradiction_clause, [], [f43256])).
fof(f43256, plain, ($false | (~ spl150_141 | ~ spl150_267 | ~ spl150_329 | ~ spl150_656 | spl150_1961)), inference(subsumption_resolution, [], [f43255, f3857])).
fof(f3857, plain, ((e22 = op2(e25, e24)) | ~ spl150_267), inference(avatar_component_clause, [], [f3855])).
fof(f3855, plain, (spl150_267 <=> (e22 = op2(e25, e24))), introduced(avatar_definition, [new_symbols(naming, [spl150_267])])).
fof(f43255, plain, (~ (e22 = op2(e25, e24)) | (~ spl150_141 | ~ spl150_329 | ~ spl150_656 | spl150_1961)), inference(forward_demodulation, [], [f43254, f6399])).
fof(f43254, plain, (~ (op2(e25, e24) = h10(e12)) | (~ spl150_141 | ~ spl150_329 | spl150_1961)), inference(forward_demodulation, [], [f43253, f3332])).
fof(f3332, plain, ((e12 = op1(e12, e10)) | ~ spl150_141), inference(avatar_component_clause, [], [f3330])).
fof(f3330, plain, (spl150_141 <=> (e12 = op1(e12, e10))), introduced(avatar_definition, [new_symbols(naming, [spl150_141])])).
fof(f43253, plain, (~ (op2(e25, e24) = h10(op1(e12, e10))) | (~ spl150_329 | spl150_1961)), inference(forward_demodulation, [], [f15223, f38500])).
fof(f38500, plain, ((e24 = h4(e13)) | ~ spl150_329), inference(backward_demodulation, [], [f1504, f4115])).
fof(f4115, plain, ((e24 = op2(e24, e20)) | ~ spl150_329), inference(avatar_component_clause, [], [f4113])).
fof(f4113, plain, (spl150_329 <=> (e24 = op2(e24, e20))), introduced(avatar_definition, [new_symbols(naming, [spl150_329])])).
fof(f1504, plain, (op2(e24, e20) = h4(e13)), inference(cnf_transformation, [], [f29])).
fof(f29, plain, ((op2(op2(e24, e20), op2(e24, e20)) = h4(e14)) & (op2(e24, e20) = h4(e13)) & (op2(e24, op2(op2(e24, e20), op2(e24, e20))) = h4(e12)) & (h4(e10) = op2(op2(e24, op2(op2(e24, e20), op2(e24, e20))), op2(e24, op2(op2(e24, e20), op2(e24, e20))))) & (e24 = h4(e15)) & (e20 = h4(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG027+1.p', ax29)).
fof(f15223, plain, (~ (h10(op1(e12, e10)) = op2(e25, h4(e13))) | spl150_1961), inference(avatar_component_clause, [], [f15221])).
fof(f15221, plain, (spl150_1961 <=> (h10(op1(e12, e10)) = op2(e25, h4(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl150_1961])])).
fof(f43215, plain, (~ spl150_137 | ~ spl150_281 | ~ spl150_969 | spl150_1960), inference(avatar_contradiction_clause, [], [f43214])).
fof(f43214, plain, ($false | (~ spl150_137 | ~ spl150_281 | ~ spl150_969 | spl150_1960)), inference(subsumption_resolution, [], [f43213, f3915])).
fof(f3915, plain, ((e24 = op2(e25, e22)) | ~ spl150_281), inference(avatar_component_clause, [], [f3913])).
fof(f3913, plain, (spl150_281 <=> (e24 = op2(e25, e22))), introduced(avatar_definition, [new_symbols(naming, [spl150_281])])).
fof(f43213, plain, (~ (e24 = op2(e25, e22)) | (~ spl150_137 | ~ spl150_969 | spl150_1960)), inference(forward_demodulation, [], [f43212, f5319])).
fof(f43212, plain, (~ (op2(e25, e22) = h10(e14)) | (~ spl150_137 | ~ spl150_969 | spl150_1960)), inference(forward_demodulation, [], [f43211, f3315])).
fof(f3315, plain, ((e14 = op1(e12, e11)) | ~ spl150_137), inference(avatar_component_clause, [], [f3313])).
fof(f3313, plain, (spl150_137 <=> (e14 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl150_137])])).
fof(f43211, plain, (~ (op2(e25, e22) = h10(op1(e12, e11))) | (~ spl150_969 | spl150_1960)), inference(forward_demodulation, [], [f15219, f8079])).
fof(f8079, plain, ((e22 = h9(e13)) | ~ spl150_969), inference(avatar_component_clause, [], [f8078])).
fof(f8078, plain, (spl150_969 <=> (e22 = h9(e13))), introduced(avatar_definition, [new_symbols(naming, [spl150_969])])).
fof(f15219, plain, (~ (h10(op1(e12, e11)) = op2(e25, h9(e13))) | spl150_1960), inference(avatar_component_clause, [], [f15217])).
fof(f15217, plain, (spl150_1960 <=> (h10(op1(e12, e11)) = op2(e25, h9(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl150_1960])])).
fof(f43178, plain, (~ spl150_116 | ~ spl150_374 | ~ spl150_656 | spl150_1958), inference(avatar_contradiction_clause, [], [f43177])).
fof(f43177, plain, ($false | (~ spl150_116 | ~ spl150_374 | ~ spl150_656 | spl150_1958)), inference(subsumption_resolution, [], [f43176, f4303])).
fof(f4303, plain, ((e21 = op2(e22, e24)) | ~ spl150_374), inference(avatar_component_clause, [], [f4301])).
fof(f4301, plain, (spl150_374 <=> (e21 = op2(e22, e24))), introduced(avatar_definition, [new_symbols(naming, [spl150_374])])).
fof(f43176, plain, (~ (e21 = op2(e22, e24)) | (~ spl150_116 | ~ spl150_656 | spl150_1958)), inference(forward_demodulation, [], [f43175, f1536])).
fof(f43175, plain, (~ (op2(e22, e24) = h10(e11)) | (~ spl150_116 | ~ spl150_656 | spl150_1958)), inference(forward_demodulation, [], [f43174, f3228])).
fof(f3228, plain, ((e11 = op1(e12, e14)) | ~ spl150_116), inference(avatar_component_clause, [], [f3226])).
fof(f3226, plain, (spl150_116 <=> (e11 = op1(e12, e14))), introduced(avatar_definition, [new_symbols(naming, [spl150_116])])).
fof(f43174, plain, (~ (op2(e22, e24) = h10(op1(e12, e14))) | (~ spl150_656 | spl150_1958)), inference(forward_demodulation, [], [f15211, f6399])).
fof(f15211, plain, (~ (h10(op1(e12, e14)) = op2(h10(e12), e24)) | spl150_1958), inference(avatar_component_clause, [], [f15209])).
fof(f15209, plain, (spl150_1958 <=> (h10(op1(e12, e14)) = op2(h10(e12), e24))), introduced(avatar_definition, [new_symbols(naming, [spl150_1958])])).
fof(f43138, plain, (~ spl150_112 | ~ spl150_286 | ~ spl150_559 | ~ spl150_954 | spl150_1957), inference(avatar_contradiction_clause, [], [f43137])).
fof(f43137, plain, ($false | (~ spl150_112 | ~ spl150_286 | ~ spl150_559 | ~ spl150_954 | spl150_1957)), inference(subsumption_resolution, [], [f43136, f3936])).
fof(f3936, plain, ((e23 = op2(e25, e21)) | ~ spl150_286), inference(avatar_component_clause, [], [f3934])).
fof(f3934, plain, (spl150_286 <=> (e23 = op2(e25, e21))), introduced(avatar_definition, [new_symbols(naming, [spl150_286])])).
fof(f43136, plain, (~ (e23 = op2(e25, e21)) | (~ spl150_112 | ~ spl150_559 | ~ spl150_954 | spl150_1957)), inference(forward_demodulation, [], [f43135, f7986])).
fof(f7986, plain, ((e23 = h10(e13)) | ~ spl150_954), inference(avatar_component_clause, [], [f7985])).
fof(f7985, plain, (spl150_954 <=> (e23 = h10(e13))), introduced(avatar_definition, [new_symbols(naming, [spl150_954])])).
fof(f43135, plain, (~ (op2(e25, e21) = h10(e13)) | (~ spl150_112 | ~ spl150_559 | spl150_1957)), inference(forward_demodulation, [], [f43134, f3211])).
fof(f3211, plain, ((e13 = op1(e12, e15)) | ~ spl150_112), inference(avatar_component_clause, [], [f3209])).
fof(f3209, plain, (spl150_112 <=> (e13 = op1(e12, e15))), introduced(avatar_definition, [new_symbols(naming, [spl150_112])])).
fof(f43134, plain, (~ (op2(e25, e21) = h10(op1(e12, e15))) | (~ spl150_559 | spl150_1957)), inference(forward_demodulation, [], [f15207, f5879])).
fof(f5879, plain, ((e21 = h30(e13)) | ~ spl150_559), inference(avatar_component_clause, [], [f5878])).
fof(f5878, plain, (spl150_559 <=> (e21 = h30(e13))), introduced(avatar_definition, [new_symbols(naming, [spl150_559])])).
fof(f15207, plain, (~ (h10(op1(e12, e15)) = op2(e25, h30(e13))) | spl150_1957), inference(avatar_component_clause, [], [f15205])).
fof(f15205, plain, (spl150_1957 <=> (h10(op1(e12, e15)) = op2(e25, h30(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl150_1957])])).
fof(f43096, plain, (~ spl150_106 | ~ spl150_286 | ~ spl150_434 | ~ spl150_954 | spl150_1956), inference(avatar_contradiction_clause, [], [f43095])).
fof(f43095, plain, ($false | (~ spl150_106 | ~ spl150_286 | ~ spl150_434 | ~ spl150_954 | spl150_1956)), inference(subsumption_resolution, [], [f43094, f3936])).
fof(f43094, plain, (~ (e23 = op2(e25, e21)) | (~ spl150_106 | ~ spl150_434 | ~ spl150_954 | spl150_1956)), inference(forward_demodulation, [], [f43093, f7986])).
fof(f43093, plain, (~ (op2(e25, e21) = h10(e13)) | (~ spl150_106 | ~ spl150_434 | spl150_1956)), inference(forward_demodulation, [], [f43092, f3186])).
fof(f3186, plain, ((e13 = op1(e13, e10)) | ~ spl150_106), inference(avatar_component_clause, [], [f3184])).
fof(f3184, plain, (spl150_106 <=> (e13 = op1(e13, e10))), introduced(avatar_definition, [new_symbols(naming, [spl150_106])])).
fof(f43092, plain, (~ (op2(e25, e21) = h10(op1(e13, e10))) | (~ spl150_434 | spl150_1956)), inference(forward_demodulation, [], [f15203, f39771])).
fof(f15203, plain, (~ (h10(op1(e13, e10)) = op2(e25, h1(e13))) | spl150_1956), inference(avatar_component_clause, [], [f15201])).
fof(f15201, plain, (spl150_1956 <=> (h10(op1(e13, e10)) = op2(e25, h1(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl150_1956])])).
fof(f43049, plain, (~ spl150_89 | spl150_1954), inference(avatar_contradiction_clause, [], [f43048])).
fof(f43048, plain, ($false | (~ spl150_89 | spl150_1954)), inference(subsumption_resolution, [], [f43047, f5319])).
fof(f43047, plain, (~ (e24 = h10(e14)) | (~ spl150_89 | spl150_1954)), inference(forward_demodulation, [], [f15195, f3115])).
fof(f3115, plain, ((e14 = op1(e13, e13)) | ~ spl150_89), inference(avatar_component_clause, [], [f3113])).
fof(f3113, plain, (spl150_89 <=> (e14 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl150_89])])).
fof(f15195, plain, (~ (e24 = h10(op1(e13, e13))) | spl150_1954), inference(avatar_component_clause, [], [f15193])).
fof(f15193, plain, (spl150_1954 <=> (e24 = h10(op1(e13, e13)))), introduced(avatar_definition, [new_symbols(naming, [spl150_1954])])).
fof(f43003, plain, (~ spl150_75 | ~ spl150_267 | ~ spl150_609 | ~ spl150_656 | spl150_1952), inference(avatar_contradiction_clause, [], [f43002])).
fof(f43002, plain, ($false | (~ spl150_75 | ~ spl150_267 | ~ spl150_609 | ~ spl150_656 | spl150_1952)), inference(subsumption_resolution, [], [f43001, f3857])).
fof(f43001, plain, (~ (e22 = op2(e25, e24)) | (~ spl150_75 | ~ spl150_609 | ~ spl150_656 | spl150_1952)), inference(forward_demodulation, [], [f43000, f6399])).
fof(f43000, plain, (~ (op2(e25, e24) = h10(e12)) | (~ spl150_75 | ~ spl150_609 | spl150_1952)), inference(forward_demodulation, [], [f42999, f3057])).
fof(f3057, plain, ((e12 = op1(e13, e15)) | ~ spl150_75), inference(avatar_component_clause, [], [f3055])).
fof(f3055, plain, (spl150_75 <=> (e12 = op1(e13, e15))), introduced(avatar_definition, [new_symbols(naming, [spl150_75])])).
fof(f42999, plain, (~ (op2(e25, e24) = h10(op1(e13, e15))) | (~ spl150_609 | spl150_1952)), inference(forward_demodulation, [], [f15187, f6141])).
fof(f15187, plain, (~ (h10(op1(e13, e15)) = op2(e25, h27(e13))) | spl150_1952), inference(avatar_component_clause, [], [f15185])).
fof(f15185, plain, (spl150_1952 <=> (h10(op1(e13, e15)) = op2(e25, h27(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl150_1952])])).
fof(f42964, plain, (~ spl150_71 | ~ spl150_329 | spl150_1951), inference(avatar_contradiction_clause, [], [f42963])).
fof(f42963, plain, ($false | (~ spl150_71 | ~ spl150_329 | spl150_1951)), inference(subsumption_resolution, [], [f42962, f38500])).
fof(f42962, plain, (~ (e24 = h4(e13)) | (~ spl150_71 | spl150_1951)), inference(forward_demodulation, [], [f42961, f5319])).
fof(f42961, plain, (~ (h4(e13) = h10(e14)) | (~ spl150_71 | spl150_1951)), inference(forward_demodulation, [], [f15183, f3040])).
fof(f3040, plain, ((e14 = op1(e14, e10)) | ~ spl150_71), inference(avatar_component_clause, [], [f3038])).
fof(f3038, plain, (spl150_71 <=> (e14 = op1(e14, e10))), introduced(avatar_definition, [new_symbols(naming, [spl150_71])])).
fof(f15183, plain, (~ (h4(e13) = h10(op1(e14, e10))) | spl150_1951), inference(avatar_component_clause, [], [f15181])).
fof(f15181, plain, (spl150_1951 <=> (h4(e13) = h10(op1(e14, e10)))), introduced(avatar_definition, [new_symbols(naming, [spl150_1951])])).
fof(f42922, plain, (~ spl150_63 | ~ spl150_656 | ~ spl150_969 | spl150_1950), inference(avatar_contradiction_clause, [], [f42921])).
fof(f42921, plain, ($false | (~ spl150_63 | ~ spl150_656 | ~ spl150_969 | spl150_1950)), inference(subsumption_resolution, [], [f42920, f8079])).
fof(f42920, plain, (~ (e22 = h9(e13)) | (~ spl150_63 | ~ spl150_656 | spl150_1950)), inference(forward_demodulation, [], [f42919, f6399])).
fof(f42919, plain, (~ (h9(e13) = h10(e12)) | (~ spl150_63 | spl150_1950)), inference(forward_demodulation, [], [f15179, f3007])).
fof(f3007, plain, ((e12 = op1(e14, e11)) | ~ spl150_63), inference(avatar_component_clause, [], [f3005])).
fof(f3005, plain, (spl150_63 <=> (e12 = op1(e14, e11))), introduced(avatar_definition, [new_symbols(naming, [spl150_63])])).
fof(f15179, plain, (~ (h9(e13) = h10(op1(e14, e11))) | spl150_1950), inference(avatar_component_clause, [], [f15177])).
fof(f15177, plain, (spl150_1950 <=> (h9(e13) = h10(op1(e14, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl150_1950])])).
fof(f42873, plain, (~ spl150_60 | ~ spl150_318 | ~ spl150_656 | spl150_1949), inference(avatar_contradiction_clause, [], [f42872])).
fof(f42872, plain, ($false | (~ spl150_60 | ~ spl150_318 | ~ spl150_656 | spl150_1949)), inference(subsumption_resolution, [], [f42871, f4069])).
fof(f4069, plain, ((e25 = op2(e24, e22)) | ~ spl150_318), inference(avatar_component_clause, [], [f4067])).
fof(f4067, plain, (spl150_318 <=> (e25 = op2(e24, e22))), introduced(avatar_definition, [new_symbols(naming, [spl150_318])])).
fof(f42871, plain, (~ (e25 = op2(e24, e22)) | (~ spl150_60 | ~ spl150_656 | spl150_1949)), inference(forward_demodulation, [], [f42870, f1537])).
fof(f42870, plain, (~ (op2(e24, e22) = h10(e15)) | (~ spl150_60 | ~ spl150_656 | spl150_1949)), inference(forward_demodulation, [], [f42869, f2994])).
fof(f2994, plain, ((e15 = op1(e14, e12)) | ~ spl150_60), inference(avatar_component_clause, [], [f2992])).
fof(f2992, plain, (spl150_60 <=> (e15 = op1(e14, e12))), introduced(avatar_definition, [new_symbols(naming, [spl150_60])])).
fof(f42869, plain, (~ (op2(e24, e22) = h10(op1(e14, e12))) | (~ spl150_656 | spl150_1949)), inference(forward_demodulation, [], [f15175, f6399])).
fof(f15175, plain, (~ (h10(op1(e14, e12)) = op2(e24, h10(e12))) | spl150_1949), inference(avatar_component_clause, [], [f15173])).
fof(f15173, plain, (spl150_1949 <=> (h10(op1(e14, e12)) = op2(e24, h10(e12)))), introduced(avatar_definition, [new_symbols(naming, [spl150_1949])])).
fof(f42826, plain, (~ spl150_38 | ~ spl150_559 | spl150_1946), inference(avatar_contradiction_clause, [], [f42825])).
fof(f42825, plain, ($false | (~ spl150_38 | ~ spl150_559 | spl150_1946)), inference(subsumption_resolution, [], [f42824, f1536])).
fof(f42824, plain, (~ (e21 = h10(e11)) | (~ spl150_38 | ~ spl150_559 | spl150_1946)), inference(forward_demodulation, [], [f42823, f5879])).
fof(f42823, plain, (~ (h10(e11) = h30(e13)) | (~ spl150_38 | spl150_1946)), inference(forward_demodulation, [], [f15163, f2903])).
fof(f2903, plain, ((e11 = op1(e14, e15)) | ~ spl150_38), inference(avatar_component_clause, [], [f2901])).
fof(f2901, plain, (spl150_38 <=> (e11 = op1(e14, e15))), introduced(avatar_definition, [new_symbols(naming, [spl150_38])])).
fof(f15163, plain, (~ (h30(e13) = h10(op1(e14, e15))) | spl150_1946), inference(avatar_component_clause, [], [f15161])).
fof(f15161, plain, (spl150_1946 <=> (h30(e13) = h10(op1(e14, e15)))), introduced(avatar_definition, [new_symbols(naming, [spl150_1946])])).
fof(f42773, plain, (~ spl150_196 | ~ spl150_454 | ~ spl150_954 | spl150_1970), inference(avatar_contradiction_clause, [], [f42772])).
fof(f42772, plain, ($false | (~ spl150_196 | ~ spl150_454 | ~ spl150_954 | spl150_1970)), inference(subsumption_resolution, [], [f42771, f4636])).
fof(f4636, plain, ((e23 = op2(e20, e23)) | ~ spl150_454), inference(avatar_component_clause, [], [f4634])).
fof(f4634, plain, (spl150_454 <=> (e23 = op2(e20, e23))), introduced(avatar_definition, [new_symbols(naming, [spl150_454])])).
fof(f42771, plain, (~ (e23 = op2(e20, e23)) | (~ spl150_196 | ~ spl150_954 | spl150_1970)), inference(forward_demodulation, [], [f42770, f7986])).
fof(f42770, plain, (~ (op2(e20, e23) = h10(e13)) | (~ spl150_196 | ~ spl150_954 | spl150_1970)), inference(forward_demodulation, [], [f42769, f3561])).
fof(f3561, plain, ((e13 = op1(e10, e13)) | ~ spl150_196), inference(avatar_component_clause, [], [f3559])).
fof(f3559, plain, (spl150_196 <=> (e13 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl150_196])])).
fof(f42769, plain, (~ (op2(e20, e23) = h10(op1(e10, e13))) | (~ spl150_954 | spl150_1970)), inference(forward_demodulation, [], [f15259, f7986])).
fof(f15259, plain, (~ (h10(op1(e10, e13)) = op2(e20, h10(e13))) | spl150_1970), inference(avatar_component_clause, [], [f15257])).
fof(f15257, plain, (spl150_1970 <=> (h10(op1(e10, e13)) = op2(e20, h10(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl150_1970])])).
fof(f42723, plain, (~ spl150_159 | ~ spl150_417 | ~ spl150_656 | ~ spl150_954 | spl150_1964), inference(avatar_contradiction_clause, [], [f42722])).
fof(f42722, plain, ($false | (~ spl150_159 | ~ spl150_417 | ~ spl150_656 | ~ spl150_954 | spl150_1964)), inference(subsumption_resolution, [], [f42721, f4482])).
fof(f4482, plain, ((e22 = op2(e21, e23)) | ~ spl150_417), inference(avatar_component_clause, [], [f4480])).
fof(f4480, plain, (spl150_417 <=> (e22 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl150_417])])).
fof(f42721, plain, (~ (e22 = op2(e21, e23)) | (~ spl150_159 | ~ spl150_656 | ~ spl150_954 | spl150_1964)), inference(forward_demodulation, [], [f42720, f6399])).
fof(f42720, plain, (~ (op2(e21, e23) = h10(e12)) | (~ spl150_159 | ~ spl150_954 | spl150_1964)), inference(forward_demodulation, [], [f42719, f3407])).
fof(f3407, plain, ((e12 = op1(e11, e13)) | ~ spl150_159), inference(avatar_component_clause, [], [f3405])).
fof(f3405, plain, (spl150_159 <=> (e12 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl150_159])])).
fof(f42719, plain, (~ (op2(e21, e23) = h10(op1(e11, e13))) | (~ spl150_954 | spl150_1964)), inference(forward_demodulation, [], [f15235, f7986])).
fof(f15235, plain, (~ (h10(op1(e11, e13)) = op2(e21, h10(e13))) | spl150_1964), inference(avatar_component_clause, [], [f15233])).
fof(f15233, plain, (spl150_1964 <=> (h10(op1(e11, e13)) = op2(e21, h10(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl150_1964])])).
fof(f42670, plain, (~ spl150_127 | spl150_1959), inference(avatar_contradiction_clause, [], [f42669])).
fof(f42669, plain, ($false | (~ spl150_127 | spl150_1959)), inference(subsumption_resolution, [], [f42668, f5365])).
fof(f42668, plain, (~ (e20 = h10(e10)) | (~ spl150_127 | spl150_1959)), inference(forward_demodulation, [], [f15215, f3274])).
fof(f3274, plain, ((e10 = op1(e12, e12)) | ~ spl150_127), inference(avatar_component_clause, [], [f3272])).
fof(f3272, plain, (spl150_127 <=> (e10 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl150_127])])).
fof(f15215, plain, (~ (e20 = h10(op1(e12, e12))) | spl150_1959), inference(avatar_component_clause, [], [f15213])).
fof(f15213, plain, (spl150_1959 <=> (e20 = h10(op1(e12, e12)))), introduced(avatar_definition, [new_symbols(naming, [spl150_1959])])).
fof(f42617, plain, (~ spl150_79 | ~ spl150_259 | ~ spl150_1495 | spl150_1953), inference(avatar_contradiction_clause, [], [f42616])).
fof(f42616, plain, ($false | (~ spl150_79 | ~ spl150_259 | ~ spl150_1495 | spl150_1953)), inference(subsumption_resolution, [], [f42615, f3824])).
fof(f3824, plain, ((e20 = op2(e25, e25)) | ~ spl150_259), inference(avatar_component_clause, [], [f3822])).
fof(f3822, plain, (spl150_259 <=> (e20 = op2(e25, e25))), introduced(avatar_definition, [new_symbols(naming, [spl150_259])])).
fof(f42615, plain, (~ (e20 = op2(e25, e25)) | (~ spl150_79 | ~ spl150_1495 | spl150_1953)), inference(forward_demodulation, [], [f42614, f5365])).
fof(f42614, plain, (~ (op2(e25, e25) = h10(e10)) | (~ spl150_79 | ~ spl150_1495 | spl150_1953)), inference(forward_demodulation, [], [f42613, f3074])).
fof(f3074, plain, ((e10 = op1(e13, e14)) | ~ spl150_79), inference(avatar_component_clause, [], [f3072])).
fof(f3072, plain, (spl150_79 <=> (e10 = op1(e13, e14))), introduced(avatar_definition, [new_symbols(naming, [spl150_79])])).
fof(f42613, plain, (~ (op2(e25, e25) = h10(op1(e13, e14))) | (~ spl150_1495 | spl150_1953)), inference(forward_demodulation, [], [f15191, f11819])).
fof(f15191, plain, (~ (h10(op1(e13, e14)) = op2(e25, h22(e13))) | spl150_1953), inference(avatar_component_clause, [], [f15189])).
fof(f15189, plain, (spl150_1953 <=> (h10(op1(e13, e14)) = op2(e25, h22(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl150_1953])])).
fof(f42562, plain, (~ spl150_49 | ~ spl150_307 | ~ spl150_954 | spl150_1948), inference(avatar_contradiction_clause, [], [f42561])).
fof(f42561, plain, ($false | (~ spl150_49 | ~ spl150_307 | ~ spl150_954 | spl150_1948)), inference(subsumption_resolution, [], [f42560, f4024])).
fof(f4024, plain, ((e20 = op2(e24, e23)) | ~ spl150_307), inference(avatar_component_clause, [], [f4022])).
fof(f4022, plain, (spl150_307 <=> (e20 = op2(e24, e23))), introduced(avatar_definition, [new_symbols(naming, [spl150_307])])).
fof(f42560, plain, (~ (e20 = op2(e24, e23)) | (~ spl150_49 | ~ spl150_954 | spl150_1948)), inference(forward_demodulation, [], [f42559, f5365])).
fof(f42559, plain, (~ (op2(e24, e23) = h10(e10)) | (~ spl150_49 | ~ spl150_954 | spl150_1948)), inference(forward_demodulation, [], [f42558, f2949])).
fof(f2949, plain, ((e10 = op1(e14, e13)) | ~ spl150_49), inference(avatar_component_clause, [], [f2947])).
fof(f2947, plain, (spl150_49 <=> (e10 = op1(e14, e13))), introduced(avatar_definition, [new_symbols(naming, [spl150_49])])).
fof(f42558, plain, (~ (op2(e24, e23) = h10(op1(e14, e13))) | (~ spl150_954 | spl150_1948)), inference(forward_demodulation, [], [f15171, f7986])).
fof(f15171, plain, (~ (h10(op1(e14, e13)) = op2(e24, h10(e13))) | spl150_1948), inference(avatar_component_clause, [], [f15169])).
fof(f15169, plain, (spl150_1948 <=> (h10(op1(e14, e13)) = op2(e24, h10(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl150_1948])])).
fof(f42502, plain, (~ spl150_36 | ~ spl150_1376 | spl150_1945), inference(avatar_contradiction_clause, [], [f42501])).
fof(f42501, plain, ($false | (~ spl150_36 | ~ spl150_1376 | spl150_1945)), inference(subsumption_resolution, [], [f42500, f10949])).
fof(f10949, plain, ((e25 = h5(e13)) | ~ spl150_1376), inference(avatar_component_clause, [], [f10948])).
fof(f10948, plain, (spl150_1376 <=> (e25 = h5(e13))), introduced(avatar_definition, [new_symbols(naming, [spl150_1376])])).
fof(f42500, plain, (~ (e25 = h5(e13)) | (~ spl150_36 | spl150_1945)), inference(forward_demodulation, [], [f42499, f1537])).
fof(f42499, plain, (~ (h5(e13) = h10(e15)) | (~ spl150_36 | spl150_1945)), inference(forward_demodulation, [], [f15159, f2894])).
fof(f2894, plain, ((e15 = op1(e15, e10)) | ~ spl150_36), inference(avatar_component_clause, [], [f2892])).
fof(f2892, plain, (spl150_36 <=> (e15 = op1(e15, e10))), introduced(avatar_definition, [new_symbols(naming, [spl150_36])])).
fof(f15159, plain, (~ (h5(e13) = h10(op1(e15, e10))) | spl150_1945), inference(avatar_component_clause, [], [f15157])).
fof(f15157, plain, (spl150_1945 <=> (h5(e13) = h10(op1(e15, e10)))), introduced(avatar_definition, [new_symbols(naming, [spl150_1945])])).
fof(f42441, plain, (~ spl150_28 | spl150_1944), inference(avatar_contradiction_clause, [], [f42440])).
fof(f42440, plain, ($false | (~ spl150_28 | spl150_1944)), inference(trivial_inequality_removal, [], [f42439])).
fof(f42439, plain, (~ (h10(e13) = h10(e13)) | (~ spl150_28 | spl150_1944)), inference(forward_demodulation, [], [f15155, f2861])).
fof(f2861, plain, ((e13 = op1(e15, e11)) | ~ spl150_28), inference(avatar_component_clause, [], [f2859])).
fof(f2859, plain, (spl150_28 <=> (e13 = op1(e15, e11))), introduced(avatar_definition, [new_symbols(naming, [spl150_28])])).
fof(f15155, plain, (~ (h10(e13) = h10(op1(e15, e11))) | spl150_1944), inference(avatar_component_clause, [], [f15153])).
fof(f15153, plain, (spl150_1944 <=> (h10(e13) = h10(op1(e15, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl150_1944])])).
fof(f42379, plain, (~ spl150_23 | ~ spl150_281 | ~ spl150_656 | spl150_1943), inference(avatar_contradiction_clause, [], [f42378])).
fof(f42378, plain, ($false | (~ spl150_23 | ~ spl150_281 | ~ spl150_656 | spl150_1943)), inference(subsumption_resolution, [], [f42377, f3915])).
fof(f42377, plain, (~ (e24 = op2(e25, e22)) | (~ spl150_23 | ~ spl150_656 | spl150_1943)), inference(forward_demodulation, [], [f42376, f5319])).
fof(f42376, plain, (~ (op2(e25, e22) = h10(e14)) | (~ spl150_23 | ~ spl150_656 | spl150_1943)), inference(forward_demodulation, [], [f42375, f2840])).
fof(f2840, plain, ((e14 = op1(e15, e12)) | ~ spl150_23), inference(avatar_component_clause, [], [f2838])).
fof(f2838, plain, (spl150_23 <=> (e14 = op1(e15, e12))), introduced(avatar_definition, [new_symbols(naming, [spl150_23])])).
fof(f42375, plain, (~ (op2(e25, e22) = h10(op1(e15, e12))) | (~ spl150_656 | spl150_1943)), inference(forward_demodulation, [], [f15151, f6399])).
fof(f15151, plain, (~ (h10(op1(e15, e12)) = op2(e25, h10(e12))) | spl150_1943), inference(avatar_component_clause, [], [f15149])).
fof(f15149, plain, (spl150_1943 <=> (h10(op1(e15, e12)) = op2(e25, h10(e12)))), introduced(avatar_definition, [new_symbols(naming, [spl150_1943])])).
fof(f42285, plain, (~ spl150_102 | ~ spl150_360 | ~ spl150_954 | spl150_1955), inference(avatar_contradiction_clause, [], [f42284])).
fof(f42284, plain, ($false | (~ spl150_102 | ~ spl150_360 | ~ spl150_954 | spl150_1955)), inference(subsumption_resolution, [], [f42283, f4244])).
fof(f4244, plain, ((e25 = op2(e23, e21)) | ~ spl150_360), inference(avatar_component_clause, [], [f4242])).
fof(f4242, plain, (spl150_360 <=> (e25 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl150_360])])).
fof(f42283, plain, (~ (e25 = op2(e23, e21)) | (~ spl150_102 | ~ spl150_954 | spl150_1955)), inference(forward_demodulation, [], [f42282, f1537])).
fof(f42282, plain, (~ (op2(e23, e21) = h10(e15)) | (~ spl150_102 | ~ spl150_954 | spl150_1955)), inference(forward_demodulation, [], [f42281, f3169])).
fof(f3169, plain, ((e15 = op1(e13, e11)) | ~ spl150_102), inference(avatar_component_clause, [], [f3167])).
fof(f3167, plain, (spl150_102 <=> (e15 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl150_102])])).
fof(f42281, plain, (~ (op2(e23, e21) = h10(op1(e13, e11))) | (~ spl150_954 | spl150_1955)), inference(forward_demodulation, [], [f15199, f7986])).
fof(f15199, plain, (~ (h10(op1(e13, e11)) = op2(h10(e13), e21)) | spl150_1955), inference(avatar_component_clause, [], [f15197])).
fof(f15197, plain, (spl150_1955 <=> (h10(op1(e13, e11)) = op2(h10(e13), e21))), introduced(avatar_definition, [new_symbols(naming, [spl150_1955])])).
fof(f42173, plain, (~ spl150_46 | ~ spl150_304 | ~ spl150_954 | spl150_1947), inference(avatar_contradiction_clause, [], [f42172])).
fof(f42172, plain, ($false | (~ spl150_46 | ~ spl150_304 | ~ spl150_954 | spl150_1947)), inference(subsumption_resolution, [], [f42171, f4011])).
fof(f4011, plain, ((e23 = op2(e24, e24)) | ~ spl150_304), inference(avatar_component_clause, [], [f4009])).
fof(f4009, plain, (spl150_304 <=> (e23 = op2(e24, e24))), introduced(avatar_definition, [new_symbols(naming, [spl150_304])])).
fof(f42171, plain, (~ (e23 = op2(e24, e24)) | (~ spl150_46 | ~ spl150_954 | spl150_1947)), inference(forward_demodulation, [], [f42170, f7986])).
fof(f42170, plain, (~ (op2(e24, e24) = h10(e13)) | (~ spl150_46 | spl150_1947)), inference(forward_demodulation, [], [f15167, f2936])).
fof(f2936, plain, ((e13 = op1(e14, e14)) | ~ spl150_46), inference(avatar_component_clause, [], [f2934])).
fof(f2934, plain, (spl150_46 <=> (e13 = op1(e14, e14))), introduced(avatar_definition, [new_symbols(naming, [spl150_46])])).
fof(f15167, plain, (~ (op2(e24, e24) = h10(op1(e14, e14))) | spl150_1947), inference(avatar_component_clause, [], [f15165])).
fof(f15165, plain, (spl150_1947 <=> (op2(e24, e24) = h10(op1(e14, e14)))), introduced(avatar_definition, [new_symbols(naming, [spl150_1947])])).
fof(f42051, plain, (~ spl150_14 | ~ spl150_272 | ~ spl150_954 | spl150_1942), inference(avatar_contradiction_clause, [], [f42050])).
fof(f42050, plain, ($false | (~ spl150_14 | ~ spl150_272 | ~ spl150_954 | spl150_1942)), inference(subsumption_resolution, [], [f42049, f3878])).
fof(f3878, plain, ((e21 = op2(e25, e23)) | ~ spl150_272), inference(avatar_component_clause, [], [f3876])).
fof(f3876, plain, (spl150_272 <=> (e21 = op2(e25, e23))), introduced(avatar_definition, [new_symbols(naming, [spl150_272])])).
fof(f42049, plain, (~ (e21 = op2(e25, e23)) | (~ spl150_14 | ~ spl150_954 | spl150_1942)), inference(forward_demodulation, [], [f42048, f1536])).
fof(f42048, plain, (~ (op2(e25, e23) = h10(e11)) | (~ spl150_14 | ~ spl150_954 | spl150_1942)), inference(forward_demodulation, [], [f42047, f2803])).
fof(f2803, plain, ((e11 = op1(e15, e13)) | ~ spl150_14), inference(avatar_component_clause, [], [f2801])).
fof(f2801, plain, (spl150_14 <=> (e11 = op1(e15, e13))), introduced(avatar_definition, [new_symbols(naming, [spl150_14])])).
fof(f42047, plain, (~ (op2(e25, e23) = h10(op1(e15, e13))) | (~ spl150_954 | spl150_1942)), inference(forward_demodulation, [], [f15147, f7986])).
fof(f15147, plain, (~ (h10(op1(e15, e13)) = op2(e25, h10(e13))) | spl150_1942), inference(avatar_component_clause, [], [f15145])).
fof(f15145, plain, (spl150_1942 <=> (h10(op1(e15, e13)) = op2(e25, h10(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl150_1942])])).
fof(f41926, plain, (~ spl150_9 | spl150_1941), inference(avatar_contradiction_clause, [], [f41925])).
fof(f41925, plain, ($false | (~ spl150_9 | spl150_1941)), inference(trivial_inequality_removal, [], [f41924])).
fof(f41924, plain, (~ (h10(e12) = h10(e12)) | (~ spl150_9 | spl150_1941)), inference(forward_demodulation, [], [f15143, f2782])).
fof(f2782, plain, ((e12 = op1(e15, e14)) | ~ spl150_9), inference(avatar_component_clause, [], [f2780])).
fof(f2780, plain, (spl150_9 <=> (e12 = op1(e15, e14))), introduced(avatar_definition, [new_symbols(naming, [spl150_9])])).
fof(f15143, plain, (~ (h10(e12) = h10(op1(e15, e14))) | spl150_1941), inference(avatar_component_clause, [], [f15141])).
fof(f15141, plain, (spl150_1941 <=> (h10(e12) = h10(op1(e15, e14)))), introduced(avatar_definition, [new_symbols(naming, [spl150_1941])])).
fof(f41817, plain, (~ spl150_1 | ~ spl150_259 | spl150_1940), inference(avatar_contradiction_clause, [], [f41816])).
fof(f41816, plain, ($false | (~ spl150_1 | ~ spl150_259 | spl150_1940)), inference(subsumption_resolution, [], [f41815, f3824])).
fof(f41815, plain, (~ (e20 = op2(e25, e25)) | (~ spl150_1 | spl150_1940)), inference(forward_demodulation, [], [f41814, f5365])).
fof(f41814, plain, (~ (op2(e25, e25) = h10(e10)) | (~ spl150_1 | spl150_1940)), inference(forward_demodulation, [], [f15139, f2749])).
fof(f2749, plain, ((e10 = op1(e15, e15)) | ~ spl150_1), inference(avatar_component_clause, [], [f2747])).
fof(f2747, plain, (spl150_1 <=> (e10 = op1(e15, e15))), introduced(avatar_definition, [new_symbols(naming, [spl150_1])])).
fof(f15139, plain, (~ (op2(e25, e25) = h10(op1(e15, e15))) | spl150_1940), inference(avatar_component_clause, [], [f15137])).
fof(f15137, plain, (spl150_1940 <=> (op2(e25, e25) = h10(op1(e15, e15)))), introduced(avatar_definition, [new_symbols(naming, [spl150_1940])])).
fof(f41716, plain, (~ spl150_166 | ~ spl150_424 | ~ spl150_656 | ~ spl150_954 | spl150_1965), inference(avatar_contradiction_clause, [], [f41715])).
fof(f41715, plain, ($false | (~ spl150_166 | ~ spl150_424 | ~ spl150_656 | ~ spl150_954 | spl150_1965)), inference(subsumption_resolution, [], [f41714, f4511])).
fof(f4511, plain, ((e23 = op2(e21, e22)) | ~ spl150_424), inference(avatar_component_clause, [], [f4509])).
fof(f4509, plain, (spl150_424 <=> (e23 = op2(e21, e22))), introduced(avatar_definition, [new_symbols(naming, [spl150_424])])).
fof(f41714, plain, (~ (e23 = op2(e21, e22)) | (~ spl150_166 | ~ spl150_656 | ~ spl150_954 | spl150_1965)), inference(forward_demodulation, [], [f41713, f7986])).
fof(f41713, plain, (~ (op2(e21, e22) = h10(e13)) | (~ spl150_166 | ~ spl150_656 | spl150_1965)), inference(forward_demodulation, [], [f41712, f3436])).
fof(f3436, plain, ((e13 = op1(e11, e12)) | ~ spl150_166), inference(avatar_component_clause, [], [f3434])).
fof(f3434, plain, (spl150_166 <=> (e13 = op1(e11, e12))), introduced(avatar_definition, [new_symbols(naming, [spl150_166])])).
fof(f41712, plain, (~ (op2(e21, e22) = h10(op1(e11, e12))) | (~ spl150_656 | spl150_1965)), inference(forward_demodulation, [], [f15239, f6399])).
fof(f15239, plain, (~ (h10(op1(e11, e12)) = op2(e21, h10(e12))) | spl150_1965), inference(avatar_component_clause, [], [f15237])).
fof(f15237, plain, (spl150_1965 <=> (h10(op1(e11, e12)) = op2(e21, h10(e12)))), introduced(avatar_definition, [new_symbols(naming, [spl150_1965])])).
fof(f41404, plain, (~ spl150_350 | ~ spl150_92 | ~ spl150_656 | ~ spl150_954 | spl150_1939), inference(avatar_split_clause, [], [f41011, f15133, f7985, f6398, f3126, f4201])).
fof(f4201, plain, (spl150_350 <=> (e21 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl150_350])])).
fof(f3126, plain, (spl150_92 <=> (e11 = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl150_92])])).
fof(f15133, plain, (spl150_1939 <=> (h10(op1(e13, e12)) = op2(h10(e13), h10(e12)))), introduced(avatar_definition, [new_symbols(naming, [spl150_1939])])).
fof(f41011, plain, (~ (e21 = op2(e23, e22)) | (~ spl150_92 | ~ spl150_656 | ~ spl150_954 | spl150_1939)), inference(forward_demodulation, [], [f41010, f1536])).
fof(f41010, plain, (~ (op2(e23, e22) = h10(e11)) | (~ spl150_92 | ~ spl150_656 | ~ spl150_954 | spl150_1939)), inference(forward_demodulation, [], [f41009, f3128])).
fof(f3128, plain, ((e11 = op1(e13, e12)) | ~ spl150_92), inference(avatar_component_clause, [], [f3126])).
fof(f41009, plain, (~ (op2(e23, e22) = h10(op1(e13, e12))) | (~ spl150_656 | ~ spl150_954 | spl150_1939)), inference(forward_demodulation, [], [f41008, f7986])).
fof(f41008, plain, (~ (h10(op1(e13, e12)) = op2(h10(e13), e22)) | (~ spl150_656 | spl150_1939)), inference(forward_demodulation, [], [f15135, f6399])).
fof(f15135, plain, (~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | spl150_1939), inference(avatar_component_clause, [], [f15133])).
fof(f41403, plain, (spl150_350 | ~ spl150_272 | ~ spl150_820 | ~ spl150_2052), inference(avatar_split_clause, [], [f41253, f15909, f7273, f3876, f4201])).
fof(f7273, plain, (spl150_820 <=> (e22 = h17(e13))), introduced(avatar_definition, [new_symbols(naming, [spl150_820])])).
fof(f15909, plain, (spl150_2052 <=> (e25 = h8(e13))), introduced(avatar_definition, [new_symbols(naming, [spl150_2052])])).
fof(f41253, plain, ((e21 = op2(e23, e22)) | (~ spl150_272 | ~ spl150_820 | ~ spl150_2052)), inference(backward_demodulation, [], [f40609, f7274])).
fof(f7274, plain, ((e22 = h17(e13)) | ~ spl150_820), inference(avatar_component_clause, [], [f7273])).
fof(f40609, plain, ((e21 = op2(e23, h17(e13))) | (~ spl150_272 | ~ spl150_2052)), inference(forward_demodulation, [], [f40608, f3878])).
fof(f40608, plain, ((op2(e25, e23) = op2(e23, h17(e13))) | ~ spl150_2052), inference(forward_demodulation, [], [f5529, f15910])).
fof(f15910, plain, ((e25 = h8(e13)) | ~ spl150_2052), inference(avatar_component_clause, [], [f15909])).
fof(f5529, plain, (op2(h8(e13), e23) = op2(e23, h17(e13))), inference(backward_demodulation, [], [f5275, f1582])).
fof(f1582, plain, (op2(e21, e23) = h17(e13)), inference(cnf_transformation, [], [f42])).
fof(f42, plain, ((op2(op2(e21, e23), op2(e21, e23)) = h17(e14)) & (op2(e21, e23) = h17(e13)) & (op2(e21, op2(op2(e21, e23), op2(e21, e23))) = h17(e12)) & (h17(e10) = op2(op2(e21, op2(op2(e21, e23), op2(e21, e23))), op2(e21, op2(op2(e21, e23), op2(e21, e23))))) & (e21 = h17(e15)) & (e23 = h17(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG027+1.p', ax42)).
fof(f5275, plain, (op2(e23, op2(e21, e23)) = op2(h8(e13), e23)), inference(backward_demodulation, [], [f800, f1528])).
fof(f1528, plain, (op2(e23, e21) = h8(e13)), inference(cnf_transformation, [], [f33])).
fof(f33, plain, ((op2(op2(e23, e21), op2(e23, e21)) = h8(e14)) & (op2(e23, e21) = h8(e13)) & (op2(e23, op2(op2(e23, e21), op2(e23, e21))) = h8(e12)) & (h8(e10) = op2(op2(e23, op2(op2(e23, e21), op2(e23, e21))), op2(e23, op2(op2(e23, e21), op2(e23, e21))))) & (e23 = h8(e15)) & (e21 = h8(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG027+1.p', ax33)).
fof(f800, plain, (op2(op2(e23, e21), e23) = op2(e23, op2(e21, e23))), inference(cnf_transformation, [], [f6])).
fof(f6, plain, ((op2(op2(e25, e25), e25) = op2(e25, op2(e25, e25))) & (op2(op2(e25, e25), e24) = op2(e25, op2(e25, e24))) & (op2(op2(e25, e25), e23) = op2(e25, op2(e25, e23))) & (op2(op2(e25, e25), e22) = op2(e25, op2(e25, e22))) & (op2(op2(e25, e25), e21) = op2(e25, op2(e25, e21))) & (op2(op2(e25, e25), e20) = op2(e25, op2(e25, e20))) & (op2(op2(e25, e24), e25) = op2(e25, op2(e24, e25))) & (op2(op2(e25, e24), e24) = op2(e25, op2(e24, e24))) & (op2(op2(e25, e24), e23) = op2(e25, op2(e24, e23))) & (op2(op2(e25, e24), e22) = op2(e25, op2(e24, e22))) & (op2(op2(e25, e24), e21) = op2(e25, op2(e24, e21))) & (op2(op2(e25, e24), e20) = op2(e25, op2(e24, e20))) & (op2(op2(e25, e23), e25) = op2(e25, op2(e23, e25))) & (op2(op2(e25, e23), e24) = op2(e25, op2(e23, e24))) & (op2(op2(e25, e23), e23) = op2(e25, op2(e23, e23))) & (op2(op2(e25, e23), e22) = op2(e25, op2(e23, e22))) & (op2(op2(e25, e23), e21) = op2(e25, op2(e23, e21))) & (op2(op2(e25, e23), e20) = op2(e25, op2(e23, e20))) & (op2(op2(e25, e22), e25) = op2(e25, op2(e22, e25))) & (op2(op2(e25, e22), e24) = op2(e25, op2(e22, e24))) & (op2(op2(e25, e22), e23) = op2(e25, op2(e22, e23))) & (op2(op2(e25, e22), e22) = op2(e25, op2(e22, e22))) & (op2(op2(e25, e22), e21) = op2(e25, op2(e22, e21))) & (op2(op2(e25, e22), e20) = op2(e25, op2(e22, e20))) & (op2(op2(e25, e21), e25) = op2(e25, op2(e21, e25))) & (op2(op2(e25, e21), e24) = op2(e25, op2(e21, e24))) & (op2(op2(e25, e21), e23) = op2(e25, op2(e21, e23))) & (op2(op2(e25, e21), e22) = op2(e25, op2(e21, e22))) & (op2(op2(e25, e21), e21) = op2(e25, op2(e21, e21))) & (op2(op2(e25, e21), e20) = op2(e25, op2(e21, e20))) & (op2(op2(e25, e20), e25) = op2(e25, op2(e20, e25))) & (op2(op2(e25, e20), e24) = op2(e25, op2(e20, e24))) & (op2(op2(e25, e20), e23) = op2(e25, op2(e20, e23))) & (op2(op2(e25, e20), e22) = op2(e25, op2(e20, e22))) & (op2(op2(e25, e20), e21) = op2(e25, op2(e20, e21))) & (op2(op2(e25, e20), e20) = op2(e25, op2(e20, e20))) & (op2(op2(e24, e25), e25) = op2(e24, op2(e25, e25))) & (op2(op2(e24, e25), e24) = op2(e24, op2(e25, e24))) & (op2(op2(e24, e25), e23) = op2(e24, op2(e25, e23))) & (op2(op2(e24, e25), e22) = op2(e24, op2(e25, e22))) & (op2(op2(e24, e25), e21) = op2(e24, op2(e25, e21))) & (op2(op2(e24, e25), e20) = op2(e24, op2(e25, e20))) & (op2(op2(e24, e24), e25) = op2(e24, op2(e24, e25))) & (op2(op2(e24, e24), e24) = op2(e24, op2(e24, e24))) & (op2(op2(e24, e24), e23) = op2(e24, op2(e24, e23))) & (op2(op2(e24, e24), e22) = op2(e24, op2(e24, e22))) & (op2(op2(e24, e24), e21) = op2(e24, op2(e24, e21))) & (op2(op2(e24, e24), e20) = op2(e24, op2(e24, e20))) & (op2(op2(e24, e23), e25) = op2(e24, op2(e23, e25))) & (op2(op2(e24, e23), e24) = op2(e24, op2(e23, e24))) & (op2(op2(e24, e23), e23) = op2(e24, op2(e23, e23))) & (op2(op2(e24, e23), e22) = op2(e24, op2(e23, e22))) & (op2(op2(e24, e23), e21) = op2(e24, op2(e23, e21))) & (op2(op2(e24, e23), e20) = op2(e24, op2(e23, e20))) & (op2(op2(e24, e22), e25) = op2(e24, op2(e22, e25))) & (op2(op2(e24, e22), e24) = op2(e24, op2(e22, e24))) & (op2(op2(e24, e22), e23) = op2(e24, op2(e22, e23))) & (op2(op2(e24, e22), e22) = op2(e24, op2(e22, e22))) & (op2(op2(e24, e22), e21) = op2(e24, op2(e22, e21))) & (op2(op2(e24, e22), e20) = op2(e24, op2(e22, e20))) & (op2(op2(e24, e21), e25) = op2(e24, op2(e21, e25))) & (op2(op2(e24, e21), e24) = op2(e24, op2(e21, e24))) & (op2(op2(e24, e21), e23) = op2(e24, op2(e21, e23))) & (op2(op2(e24, e21), e22) = op2(e24, op2(e21, e22))) & (op2(op2(e24, e21), e21) = op2(e24, op2(e21, e21))) & (op2(op2(e24, e21), e20) = op2(e24, op2(e21, e20))) & (op2(op2(e24, e20), e25) = op2(e24, op2(e20, e25))) & (op2(op2(e24, e20), e24) = op2(e24, op2(e20, e24))) & (op2(op2(e24, e20), e23) = op2(e24, op2(e20, e23))) & (op2(op2(e24, e20), e22) = op2(e24, op2(e20, e22))) & (op2(op2(e24, e20), e21) = op2(e24, op2(e20, e21))) & (op2(op2(e24, e20), e20) = op2(e24, op2(e20, e20))) & (op2(op2(e23, e25), e25) = op2(e23, op2(e25, e25))) & (op2(op2(e23, e25), e24) = op2(e23, op2(e25, e24))) & (op2(op2(e23, e25), e23) = op2(e23, op2(e25, e23))) & (op2(op2(e23, e25), e22) = op2(e23, op2(e25, e22))) & (op2(op2(e23, e25), e21) = op2(e23, op2(e25, e21))) & (op2(op2(e23, e25), e20) = op2(e23, op2(e25, e20))) & (op2(op2(e23, e24), e25) = op2(e23, op2(e24, e25))) & (op2(op2(e23, e24), e24) = op2(e23, op2(e24, e24))) & (op2(op2(e23, e24), e23) = op2(e23, op2(e24, e23))) & (op2(op2(e23, e24), e22) = op2(e23, op2(e24, e22))) & (op2(op2(e23, e24), e21) = op2(e23, op2(e24, e21))) & (op2(op2(e23, e24), e20) = op2(e23, op2(e24, e20))) & (op2(op2(e23, e23), e25) = op2(e23, op2(e23, e25))) & (op2(op2(e23, e23), e24) = op2(e23, op2(e23, e24))) & (op2(op2(e23, e23), e23) = op2(e23, op2(e23, e23))) & (op2(op2(e23, e23), e22) = op2(e23, op2(e23, e22))) & (op2(op2(e23, e23), e21) = op2(e23, op2(e23, e21))) & (op2(op2(e23, e23), e20) = op2(e23, op2(e23, e20))) & (op2(op2(e23, e22), e25) = op2(e23, op2(e22, e25))) & (op2(op2(e23, e22), e24) = op2(e23, op2(e22, e24))) & (op2(op2(e23, e22), e23) = op2(e23, op2(e22, e23))) & (op2(op2(e23, e22), e22) = op2(e23, op2(e22, e22))) & (op2(op2(e23, e22), e21) = op2(e23, op2(e22, e21))) & (op2(op2(e23, e22), e20) = op2(e23, op2(e22, e20))) & (op2(op2(e23, e21), e25) = op2(e23, op2(e21, e25))) & (op2(op2(e23, e21), e24) = op2(e23, op2(e21, e24))) & (op2(op2(e23, e21), e23) = op2(e23, op2(e21, e23))) & (op2(op2(e23, e21), e22) = op2(e23, op2(e21, e22))) & (op2(op2(e23, e21), e21) = op2(e23, op2(e21, e21))) & (op2(op2(e23, e21), e20) = op2(e23, op2(e21, e20))) & (op2(op2(e23, e20), e25) = op2(e23, op2(e20, e25))) & (op2(op2(e23, e20), e24) = op2(e23, op2(e20, e24))) & (op2(op2(e23, e20), e23) = op2(e23, op2(e20, e23))) & (op2(op2(e23, e20), e22) = op2(e23, op2(e20, e22))) & (op2(op2(e23, e20), e21) = op2(e23, op2(e20, e21))) & (op2(op2(e23, e20), e20) = op2(e23, op2(e20, e20))) & (op2(op2(e22, e25), e25) = op2(e22, op2(e25, e25))) & (op2(op2(e22, e25), e24) = op2(e22, op2(e25, e24))) & (op2(op2(e22, e25), e23) = op2(e22, op2(e25, e23))) & (op2(op2(e22, e25), e22) = op2(e22, op2(e25, e22))) & (op2(op2(e22, e25), e21) = op2(e22, op2(e25, e21))) & (op2(op2(e22, e25), e20) = op2(e22, op2(e25, e20))) & (op2(op2(e22, e24), e25) = op2(e22, op2(e24, e25))) & (op2(op2(e22, e24), e24) = op2(e22, op2(e24, e24))) & (op2(op2(e22, e24), e23) = op2(e22, op2(e24, e23))) & (op2(op2(e22, e24), e22) = op2(e22, op2(e24, e22))) & (op2(op2(e22, e24), e21) = op2(e22, op2(e24, e21))) & (op2(op2(e22, e24), e20) = op2(e22, op2(e24, e20))) & (op2(op2(e22, e23), e25) = op2(e22, op2(e23, e25))) & (op2(op2(e22, e23), e24) = op2(e22, op2(e23, e24))) & (op2(op2(e22, e23), e23) = op2(e22, op2(e23, e23))) & (op2(op2(e22, e23), e22) = op2(e22, op2(e23, e22))) & (op2(op2(e22, e23), e21) = op2(e22, op2(e23, e21))) & (op2(op2(e22, e23), e20) = op2(e22, op2(e23, e20))) & (op2(op2(e22, e22), e25) = op2(e22, op2(e22, e25))) & (op2(op2(e22, e22), e24) = op2(e22, op2(e22, e24))) & (op2(op2(e22, e22), e23) = op2(e22, op2(e22, e23))) & (op2(op2(e22, e22), e22) = op2(e22, op2(e22, e22))) & (op2(op2(e22, e22), e21) = op2(e22, op2(e22, e21))) & (op2(op2(e22, e22), e20) = op2(e22, op2(e22, e20))) & (op2(op2(e22, e21), e25) = op2(e22, op2(e21, e25))) & (op2(op2(e22, e21), e24) = op2(e22, op2(e21, e24))) & (op2(op2(e22, e21), e23) = op2(e22, op2(e21, e23))) & (op2(op2(e22, e21), e22) = op2(e22, op2(e21, e22))) & (op2(op2(e22, e21), e21) = op2(e22, op2(e21, e21))) & (op2(op2(e22, e21), e20) = op2(e22, op2(e21, e20))) & (op2(op2(e22, e20), e25) = op2(e22, op2(e20, e25))) & (op2(op2(e22, e20), e24) = op2(e22, op2(e20, e24))) & (op2(op2(e22, e20), e23) = op2(e22, op2(e20, e23))) & (op2(op2(e22, e20), e22) = op2(e22, op2(e20, e22))) & (op2(op2(e22, e20), e21) = op2(e22, op2(e20, e21))) & (op2(op2(e22, e20), e20) = op2(e22, op2(e20, e20))) & (op2(op2(e21, e25), e25) = op2(e21, op2(e25, e25))) & (op2(op2(e21, e25), e24) = op2(e21, op2(e25, e24))) & (op2(op2(e21, e25), e23) = op2(e21, op2(e25, e23))) & (op2(op2(e21, e25), e22) = op2(e21, op2(e25, e22))) & (op2(op2(e21, e25), e21) = op2(e21, op2(e25, e21))) & (op2(op2(e21, e25), e20) = op2(e21, op2(e25, e20))) & (op2(op2(e21, e24), e25) = op2(e21, op2(e24, e25))) & (op2(op2(e21, e24), e24) = op2(e21, op2(e24, e24))) & (op2(op2(e21, e24), e23) = op2(e21, op2(e24, e23))) & (op2(op2(e21, e24), e22) = op2(e21, op2(e24, e22))) & (op2(op2(e21, e24), e21) = op2(e21, op2(e24, e21))) & (op2(op2(e21, e24), e20) = op2(e21, op2(e24, e20))) & (op2(op2(e21, e23), e25) = op2(e21, op2(e23, e25))) & (op2(op2(e21, e23), e24) = op2(e21, op2(e23, e24))) & (op2(op2(e21, e23), e23) = op2(e21, op2(e23, e23))) & (op2(op2(e21, e23), e22) = op2(e21, op2(e23, e22))) & (op2(op2(e21, e23), e21) = op2(e21, op2(e23, e21))) & (op2(op2(e21, e23), e20) = op2(e21, op2(e23, e20))) & (op2(op2(e21, e22), e25) = op2(e21, op2(e22, e25))) & (op2(op2(e21, e22), e24) = op2(e21, op2(e22, e24))) & (op2(op2(e21, e22), e23) = op2(e21, op2(e22, e23))) & (op2(op2(e21, e22), e22) = op2(e21, op2(e22, e22))) & (op2(op2(e21, e22), e21) = op2(e21, op2(e22, e21))) & (op2(op2(e21, e22), e20) = op2(e21, op2(e22, e20))) & (op2(op2(e21, e21), e25) = op2(e21, op2(e21, e25))) & (op2(op2(e21, e21), e24) = op2(e21, op2(e21, e24))) & (op2(op2(e21, e21), e23) = op2(e21, op2(e21, e23))) & (op2(op2(e21, e21), e22) = op2(e21, op2(e21, e22))) & (op2(op2(e21, e21), e21) = op2(e21, op2(e21, e21))) & (op2(op2(e21, e21), e20) = op2(e21, op2(e21, e20))) & (op2(op2(e21, e20), e25) = op2(e21, op2(e20, e25))) & (op2(op2(e21, e20), e24) = op2(e21, op2(e20, e24))) & (op2(op2(e21, e20), e23) = op2(e21, op2(e20, e23))) & (op2(op2(e21, e20), e22) = op2(e21, op2(e20, e22))) & (op2(op2(e21, e20), e21) = op2(e21, op2(e20, e21))) & (op2(op2(e21, e20), e20) = op2(e21, op2(e20, e20))) & (op2(op2(e20, e25), e25) = op2(e20, op2(e25, e25))) & (op2(op2(e20, e25), e24) = op2(e20, op2(e25, e24))) & (op2(op2(e20, e25), e23) = op2(e20, op2(e25, e23))) & (op2(op2(e20, e25), e22) = op2(e20, op2(e25, e22))) & (op2(op2(e20, e25), e21) = op2(e20, op2(e25, e21))) & (op2(op2(e20, e25), e20) = op2(e20, op2(e25, e20))) & (op2(op2(e20, e24), e25) = op2(e20, op2(e24, e25))) & (op2(op2(e20, e24), e24) = op2(e20, op2(e24, e24))) & (op2(op2(e20, e24), e23) = op2(e20, op2(e24, e23))) & (op2(op2(e20, e24), e22) = op2(e20, op2(e24, e22))) & (op2(op2(e20, e24), e21) = op2(e20, op2(e24, e21))) & (op2(op2(e20, e24), e20) = op2(e20, op2(e24, e20))) & (op2(op2(e20, e23), e25) = op2(e20, op2(e23, e25))) & (op2(op2(e20, e23), e24) = op2(e20, op2(e23, e24))) & (op2(op2(e20, e23), e23) = op2(e20, op2(e23, e23))) & (op2(op2(e20, e23), e22) = op2(e20, op2(e23, e22))) & (op2(op2(e20, e23), e21) = op2(e20, op2(e23, e21))) & (op2(op2(e20, e23), e20) = op2(e20, op2(e23, e20))) & (op2(op2(e20, e22), e25) = op2(e20, op2(e22, e25))) & (op2(op2(e20, e22), e24) = op2(e20, op2(e22, e24))) & (op2(op2(e20, e22), e23) = op2(e20, op2(e22, e23))) & (op2(op2(e20, e22), e22) = op2(e20, op2(e22, e22))) & (op2(op2(e20, e22), e21) = op2(e20, op2(e22, e21))) & (op2(op2(e20, e22), e20) = op2(e20, op2(e22, e20))) & (op2(op2(e20, e21), e25) = op2(e20, op2(e21, e25))) & (op2(op2(e20, e21), e24) = op2(e20, op2(e21, e24))) & (op2(op2(e20, e21), e23) = op2(e20, op2(e21, e23))) & (op2(op2(e20, e21), e22) = op2(e20, op2(e21, e22))) & (op2(op2(e20, e21), e21) = op2(e20, op2(e21, e21))) & (op2(op2(e20, e21), e20) = op2(e20, op2(e21, e20))) & (op2(op2(e20, e20), e25) = op2(e20, op2(e20, e25))) & (op2(op2(e20, e20), e24) = op2(e20, op2(e20, e24))) & (op2(op2(e20, e20), e23) = op2(e20, op2(e20, e23))) & (op2(op2(e20, e20), e22) = op2(e20, op2(e20, e22))) & (op2(op2(e20, e20), e21) = op2(e20, op2(e20, e21))) & (op2(op2(e20, e20), e20) = op2(e20, op2(e20, e20)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG027+1.p', ax6)).
fof(f41387, plain, (spl150_25 | ~ spl150_217 | ~ spl150_252), inference(avatar_contradiction_clause, [], [f41386])).
fof(f41386, plain, ($false | (spl150_25 | ~ spl150_217 | ~ spl150_252)), inference(subsumption_resolution, [], [f41383, f2848])).
fof(f2848, plain, (~ (e10 = op1(e15, e11)) | spl150_25), inference(avatar_component_clause, [], [f2847])).
fof(f2847, plain, (spl150_25 <=> (e10 = op1(e15, e11))), introduced(avatar_definition, [new_symbols(naming, [spl150_25])])).
fof(f41383, plain, ((e10 = op1(e15, e11)) | (~ spl150_217 | ~ spl150_252)), inference(backward_demodulation, [], [f37273, f3794])).
fof(f3794, plain, ((e15 = inv1(e11)) | ~ spl150_252), inference(avatar_component_clause, [], [f3792])).
fof(f3792, plain, (spl150_252 <=> (e15 = inv1(e11))), introduced(avatar_definition, [new_symbols(naming, [spl150_252])])).
fof(f37273, plain, ((e10 = op1(inv1(e11), e11)) | ~ spl150_217), inference(forward_demodulation, [], [f632, f3649])).
fof(f3649, plain, ((e10 = unit1) | ~ spl150_217), inference(avatar_component_clause, [], [f3647])).
fof(f3647, plain, (spl150_217 <=> (e10 = unit1)), introduced(avatar_definition, [new_symbols(naming, [spl150_217])])).
fof(f632, plain, (unit1 = op1(inv1(e11), e11)), inference(cnf_transformation, [], [f4])).
fof(f4, plain, (((e15 = inv1(e15)) | (e14 = inv1(e15)) | (e13 = inv1(e15)) | (e12 = inv1(e15)) | (e11 = inv1(e15)) | (e10 = inv1(e15))) & ((e15 = inv1(e14)) | (e14 = inv1(e14)) | (e13 = inv1(e14)) | (e12 = inv1(e14)) | (e11 = inv1(e14)) | (e10 = inv1(e14))) & ((e15 = inv1(e13)) | (e14 = inv1(e13)) | (e13 = inv1(e13)) | (e12 = inv1(e13)) | (e11 = inv1(e13)) | (e10 = inv1(e13))) & ((e15 = inv1(e12)) | (e14 = inv1(e12)) | (e13 = inv1(e12)) | (e12 = inv1(e12)) | (e11 = inv1(e12)) | (e10 = inv1(e12))) & ((e15 = inv1(e11)) | (e14 = inv1(e11)) | (e13 = inv1(e11)) | (e12 = inv1(e11)) | (e11 = inv1(e11)) | (e10 = inv1(e11))) & ((e15 = inv1(e10)) | (e14 = inv1(e10)) | (e13 = inv1(e10)) | (e12 = inv1(e10)) | (e11 = inv1(e10)) | (e10 = inv1(e10))) & (unit1 = op1(inv1(e15), e15)) & (unit1 = op1(e15, inv1(e15))) & (unit1 = op1(inv1(e14), e14)) & (unit1 = op1(e14, inv1(e14))) & (unit1 = op1(inv1(e13), e13)) & (unit1 = op1(e13, inv1(e13))) & (unit1 = op1(inv1(e12), e12)) & (unit1 = op1(e12, inv1(e12))) & (unit1 = op1(inv1(e11), e11)) & (unit1 = op1(e11, inv1(e11))) & (unit1 = op1(inv1(e10), e10)) & (unit1 = op1(e10, inv1(e10)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG027+1.p', ax4)).
fof(f41271, plain, (spl150_424 | ~ spl150_427 | ~ spl150_454 | ~ spl150_820), inference(avatar_split_clause, [], [f41256, f7273, f4634, f4522, f4509])).
fof(f41256, plain, ((e23 = op2(e21, e22)) | (~ spl150_427 | ~ spl150_454 | ~ spl150_820)), inference(backward_demodulation, [], [f40680, f7274])).
fof(f40680, plain, ((e23 = op2(e21, h17(e13))) | (~ spl150_427 | ~ spl150_454)), inference(forward_demodulation, [], [f40679, f4636])).
fof(f40679, plain, ((op2(e20, e23) = op2(e21, h17(e13))) | ~ spl150_427), inference(forward_demodulation, [], [f5513, f4524])).
fof(f5513, plain, (op2(op2(e21, e21), e23) = op2(e21, h17(e13))), inference(backward_demodulation, [], [f728, f1582])).
fof(f728, plain, (op2(op2(e21, e21), e23) = op2(e21, op2(e21, e23))), inference(cnf_transformation, [], [f6])).
fof(f40868, plain, (spl150_304 | ~ spl150_347 | ~ spl150_454 | ~ spl150_785), inference(avatar_split_clause, [], [f40867, f7088, f4634, f4188, f4009])).
fof(f4188, plain, (spl150_347 <=> (e24 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl150_347])])).
fof(f7088, plain, (spl150_785 <=> (e20 = h19(e13))), introduced(avatar_definition, [new_symbols(naming, [spl150_785])])).
fof(f40867, plain, ((e23 = op2(e24, e24)) | (~ spl150_347 | ~ spl150_454 | ~ spl150_785)), inference(forward_demodulation, [], [f40866, f4636])).
fof(f40866, plain, ((op2(e20, e23) = op2(e24, e24)) | (~ spl150_347 | ~ spl150_785)), inference(forward_demodulation, [], [f29139, f7089])).
fof(f7089, plain, ((e20 = h19(e13)) | ~ spl150_785), inference(avatar_component_clause, [], [f7088])).
fof(f29139, plain, ((op2(e24, e24) = op2(h19(e13), e23)) | ~ spl150_347), inference(forward_demodulation, [], [f5566, f4190])).
fof(f4190, plain, ((e24 = op2(e23, e23)) | ~ spl150_347), inference(avatar_component_clause, [], [f4188])).
fof(f5566, plain, (op2(e24, op2(e23, e23)) = op2(h19(e13), e23)), inference(backward_demodulation, [], [f848, f1594])).
fof(f1594, plain, (op2(e24, e23) = h19(e13)), inference(cnf_transformation, [], [f44])).
fof(f44, plain, ((op2(op2(e24, e23), op2(e24, e23)) = h19(e14)) & (op2(e24, e23) = h19(e13)) & (op2(e24, op2(op2(e24, e23), op2(e24, e23))) = h19(e12)) & (h19(e10) = op2(op2(e24, op2(op2(e24, e23), op2(e24, e23))), op2(e24, op2(op2(e24, e23), op2(e24, e23))))) & (e24 = h19(e15)) & (e23 = h19(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG027+1.p', ax44)).
fof(f848, plain, (op2(op2(e24, e23), e23) = op2(e24, op2(e23, e23))), inference(cnf_transformation, [], [f6])).
fof(f40865, plain, (~ spl150_301 | ~ spl150_785), inference(avatar_split_clause, [], [f40864, f7088, f3997])).
fof(f3997, plain, (spl150_301 <=> (e20 = op2(e24, e24))), introduced(avatar_definition, [new_symbols(naming, [spl150_301])])).
fof(f40864, plain, (~ (e20 = op2(e24, e24)) | ~ spl150_785), inference(forward_demodulation, [], [f5572, f7089])).
fof(f5572, plain, ~ (op2(e24, e24) = h19(e13)), inference(backward_demodulation, [], [f1388, f1594])).
fof(f1388, plain, ~ (op2(e24, e23) = op2(e24, e24)), inference(cnf_transformation, [], [f18])).
fof(f18, plain, (~ (op2(e25, e24) = op2(e25, e25)) & ~ (op2(e25, e23) = op2(e25, e25)) & ~ (op2(e25, e23) = op2(e25, e24)) & ~ (op2(e25, e22) = op2(e25, e25)) & ~ (op2(e25, e22) = op2(e25, e24)) & ~ (op2(e25, e22) = op2(e25, e23)) & ~ (op2(e25, e21) = op2(e25, e25)) & ~ (op2(e25, e21) = op2(e25, e24)) & ~ (op2(e25, e21) = op2(e25, e23)) & ~ (op2(e25, e21) = op2(e25, e22)) & ~ (op2(e25, e20) = op2(e25, e25)) & ~ (op2(e25, e20) = op2(e25, e24)) & ~ (op2(e25, e20) = op2(e25, e23)) & ~ (op2(e25, e20) = op2(e25, e22)) & ~ (op2(e25, e20) = op2(e25, e21)) & ~ (op2(e24, e24) = op2(e24, e25)) & ~ (op2(e24, e23) = op2(e24, e25)) & ~ (op2(e24, e23) = op2(e24, e24)) & ~ (op2(e24, e22) = op2(e24, e25)) & ~ (op2(e24, e22) = op2(e24, e24)) & ~ (op2(e24, e22) = op2(e24, e23)) & ~ (op2(e24, e21) = op2(e24, e25)) & ~ (op2(e24, e21) = op2(e24, e24)) & ~ (op2(e24, e21) = op2(e24, e23)) & ~ (op2(e24, e21) = op2(e24, e22)) & ~ (op2(e24, e20) = op2(e24, e25)) & ~ (op2(e24, e20) = op2(e24, e24)) & ~ (op2(e24, e20) = op2(e24, e23)) & ~ (op2(e24, e20) = op2(e24, e22)) & ~ (op2(e24, e20) = op2(e24, e21)) & ~ (op2(e23, e24) = op2(e23, e25)) & ~ (op2(e23, e23) = op2(e23, e25)) & ~ (op2(e23, e23) = op2(e23, e24)) & ~ (op2(e23, e22) = op2(e23, e25)) & ~ (op2(e23, e22) = op2(e23, e24)) & ~ (op2(e23, e22) = op2(e23, e23)) & ~ (op2(e23, e21) = op2(e23, e25)) & ~ (op2(e23, e21) = op2(e23, e24)) & ~ (op2(e23, e21) = op2(e23, e23)) & ~ (op2(e23, e21) = op2(e23, e22)) & ~ (op2(e23, e20) = op2(e23, e25)) & ~ (op2(e23, e20) = op2(e23, e24)) & ~ (op2(e23, e20) = op2(e23, e23)) & ~ (op2(e23, e20) = op2(e23, e22)) & ~ (op2(e23, e20) = op2(e23, e21)) & ~ (op2(e22, e24) = op2(e22, e25)) & ~ (op2(e22, e23) = op2(e22, e25)) & ~ (op2(e22, e23) = op2(e22, e24)) & ~ (op2(e22, e22) = op2(e22, e25)) & ~ (op2(e22, e22) = op2(e22, e24)) & ~ (op2(e22, e22) = op2(e22, e23)) & ~ (op2(e22, e21) = op2(e22, e25)) & ~ (op2(e22, e21) = op2(e22, e24)) & ~ (op2(e22, e21) = op2(e22, e23)) & ~ (op2(e22, e21) = op2(e22, e22)) & ~ (op2(e22, e20) = op2(e22, e25)) & ~ (op2(e22, e20) = op2(e22, e24)) & ~ (op2(e22, e20) = op2(e22, e23)) & ~ (op2(e22, e20) = op2(e22, e22)) & ~ (op2(e22, e20) = op2(e22, e21)) & ~ (op2(e21, e24) = op2(e21, e25)) & ~ (op2(e21, e23) = op2(e21, e25)) & ~ (op2(e21, e23) = op2(e21, e24)) & ~ (op2(e21, e22) = op2(e21, e25)) & ~ (op2(e21, e22) = op2(e21, e24)) & ~ (op2(e21, e22) = op2(e21, e23)) & ~ (op2(e21, e21) = op2(e21, e25)) & ~ (op2(e21, e21) = op2(e21, e24)) & ~ (op2(e21, e21) = op2(e21, e23)) & ~ (op2(e21, e21) = op2(e21, e22)) & ~ (op2(e21, e20) = op2(e21, e25)) & ~ (op2(e21, e20) = op2(e21, e24)) & ~ (op2(e21, e20) = op2(e21, e23)) & ~ (op2(e21, e20) = op2(e21, e22)) & ~ (op2(e21, e20) = op2(e21, e21)) & ~ (op2(e20, e24) = op2(e20, e25)) & ~ (op2(e20, e23) = op2(e20, e25)) & ~ (op2(e20, e23) = op2(e20, e24)) & ~ (op2(e20, e22) = op2(e20, e25)) & ~ (op2(e20, e22) = op2(e20, e24)) & ~ (op2(e20, e22) = op2(e20, e23)) & ~ (op2(e20, e21) = op2(e20, e25)) & ~ (op2(e20, e21) = op2(e20, e24)) & ~ (op2(e20, e21) = op2(e20, e23)) & ~ (op2(e20, e21) = op2(e20, e22)) & ~ (op2(e20, e20) = op2(e20, e25)) & ~ (op2(e20, e20) = op2(e20, e24)) & ~ (op2(e20, e20) = op2(e20, e23)) & ~ (op2(e20, e20) = op2(e20, e22)) & ~ (op2(e20, e20) = op2(e20, e21)) & ~ (op2(e24, e25) = op2(e25, e25)) & ~ (op2(e23, e25) = op2(e25, e25)) & ~ (op2(e23, e25) = op2(e24, e25)) & ~ (op2(e22, e25) = op2(e25, e25)) & ~ (op2(e22, e25) = op2(e24, e25)) & ~ (op2(e22, e25) = op2(e23, e25)) & ~ (op2(e21, e25) = op2(e25, e25)) & ~ (op2(e21, e25) = op2(e24, e25)) & ~ (op2(e21, e25) = op2(e23, e25)) & ~ (op2(e21, e25) = op2(e22, e25)) & ~ (op2(e20, e25) = op2(e25, e25)) & ~ (op2(e20, e25) = op2(e24, e25)) & ~ (op2(e20, e25) = op2(e23, e25)) & ~ (op2(e20, e25) = op2(e22, e25)) & ~ (op2(e20, e25) = op2(e21, e25)) & ~ (op2(e24, e24) = op2(e25, e24)) & ~ (op2(e23, e24) = op2(e25, e24)) & ~ (op2(e23, e24) = op2(e24, e24)) & ~ (op2(e22, e24) = op2(e25, e24)) & ~ (op2(e22, e24) = op2(e24, e24)) & ~ (op2(e22, e24) = op2(e23, e24)) & ~ (op2(e21, e24) = op2(e25, e24)) & ~ (op2(e21, e24) = op2(e24, e24)) & ~ (op2(e21, e24) = op2(e23, e24)) & ~ (op2(e21, e24) = op2(e22, e24)) & ~ (op2(e20, e24) = op2(e25, e24)) & ~ (op2(e20, e24) = op2(e24, e24)) & ~ (op2(e20, e24) = op2(e23, e24)) & ~ (op2(e20, e24) = op2(e22, e24)) & ~ (op2(e20, e24) = op2(e21, e24)) & ~ (op2(e24, e23) = op2(e25, e23)) & ~ (op2(e23, e23) = op2(e25, e23)) & ~ (op2(e23, e23) = op2(e24, e23)) & ~ (op2(e22, e23) = op2(e25, e23)) & ~ (op2(e22, e23) = op2(e24, e23)) & ~ (op2(e22, e23) = op2(e23, e23)) & ~ (op2(e21, e23) = op2(e25, e23)) & ~ (op2(e21, e23) = op2(e24, e23)) & ~ (op2(e21, e23) = op2(e23, e23)) & ~ (op2(e21, e23) = op2(e22, e23)) & ~ (op2(e20, e23) = op2(e25, e23)) & ~ (op2(e20, e23) = op2(e24, e23)) & ~ (op2(e20, e23) = op2(e23, e23)) & ~ (op2(e20, e23) = op2(e22, e23)) & ~ (op2(e20, e23) = op2(e21, e23)) & ~ (op2(e24, e22) = op2(e25, e22)) & ~ (op2(e23, e22) = op2(e25, e22)) & ~ (op2(e23, e22) = op2(e24, e22)) & ~ (op2(e22, e22) = op2(e25, e22)) & ~ (op2(e22, e22) = op2(e24, e22)) & ~ (op2(e22, e22) = op2(e23, e22)) & ~ (op2(e21, e22) = op2(e25, e22)) & ~ (op2(e21, e22) = op2(e24, e22)) & ~ (op2(e21, e22) = op2(e23, e22)) & ~ (op2(e21, e22) = op2(e22, e22)) & ~ (op2(e20, e22) = op2(e25, e22)) & ~ (op2(e20, e22) = op2(e24, e22)) & ~ (op2(e20, e22) = op2(e23, e22)) & ~ (op2(e20, e22) = op2(e22, e22)) & ~ (op2(e20, e22) = op2(e21, e22)) & ~ (op2(e24, e21) = op2(e25, e21)) & ~ (op2(e23, e21) = op2(e25, e21)) & ~ (op2(e23, e21) = op2(e24, e21)) & ~ (op2(e22, e21) = op2(e25, e21)) & ~ (op2(e22, e21) = op2(e24, e21)) & ~ (op2(e22, e21) = op2(e23, e21)) & ~ (op2(e21, e21) = op2(e25, e21)) & ~ (op2(e21, e21) = op2(e24, e21)) & ~ (op2(e21, e21) = op2(e23, e21)) & ~ (op2(e21, e21) = op2(e22, e21)) & ~ (op2(e20, e21) = op2(e25, e21)) & ~ (op2(e20, e21) = op2(e24, e21)) & ~ (op2(e20, e21) = op2(e23, e21)) & ~ (op2(e20, e21) = op2(e22, e21)) & ~ (op2(e20, e21) = op2(e21, e21)) & ~ (op2(e24, e20) = op2(e25, e20)) & ~ (op2(e23, e20) = op2(e25, e20)) & ~ (op2(e23, e20) = op2(e24, e20)) & ~ (op2(e22, e20) = op2(e25, e20)) & ~ (op2(e22, e20) = op2(e24, e20)) & ~ (op2(e22, e20) = op2(e23, e20)) & ~ (op2(e21, e20) = op2(e25, e20)) & ~ (op2(e21, e20) = op2(e24, e20)) & ~ (op2(e21, e20) = op2(e23, e20)) & ~ (op2(e21, e20) = op2(e22, e20)) & ~ (op2(e20, e20) = op2(e25, e20)) & ~ (op2(e20, e20) = op2(e24, e20)) & ~ (op2(e20, e20) = op2(e23, e20)) & ~ (op2(e20, e20) = op2(e22, e20)) & ~ (op2(e20, e20) = op2(e21, e20))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG027+1.p', ax18)).
fof(f40743, plain, (spl150_374 | ~ spl150_272 | ~ spl150_347 | ~ spl150_1655), inference(avatar_split_clause, [], [f40742, f12977, f4188, f3876, f4301])).
fof(f12977, plain, (spl150_1655 <=> (e25 = h18(e13))), introduced(avatar_definition, [new_symbols(naming, [spl150_1655])])).
fof(f40742, plain, ((e21 = op2(e22, e24)) | (~ spl150_272 | ~ spl150_347 | ~ spl150_1655)), inference(forward_demodulation, [], [f40741, f3878])).
fof(f40741, plain, ((op2(e22, e24) = op2(e25, e23)) | (~ spl150_347 | ~ spl150_1655)), inference(forward_demodulation, [], [f29187, f12978])).
fof(f12978, plain, ((e25 = h18(e13)) | ~ spl150_1655), inference(avatar_component_clause, [], [f12977])).
fof(f29187, plain, ((op2(e22, e24) = op2(h18(e13), e23)) | ~ spl150_347), inference(forward_demodulation, [], [f5539, f4190])).
fof(f5539, plain, (op2(e22, op2(e23, e23)) = op2(h18(e13), e23)), inference(backward_demodulation, [], [f776, f1588])).
fof(f1588, plain, (op2(e22, e23) = h18(e13)), inference(cnf_transformation, [], [f43])).
fof(f43, plain, ((op2(op2(e22, e23), op2(e22, e23)) = h18(e14)) & (op2(e22, e23) = h18(e13)) & (op2(e22, op2(op2(e22, e23), op2(e22, e23))) = h18(e12)) & (h18(e10) = op2(op2(e22, op2(op2(e22, e23), op2(e22, e23))), op2(e22, op2(op2(e22, e23), op2(e22, e23))))) & (e22 = h18(e15)) & (e23 = h18(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG027+1.p', ax43)).
fof(f776, plain, (op2(op2(e22, e23), e23) = op2(e22, op2(e23, e23))), inference(cnf_transformation, [], [f6])).
fof(f40660, plain, (spl150_417 | ~ spl150_267 | ~ spl150_347 | ~ spl150_760), inference(avatar_split_clause, [], [f39478, f6956, f4188, f3855, f4480])).
fof(f6956, plain, (spl150_760 <=> (e21 = h20(e13))), introduced(avatar_definition, [new_symbols(naming, [spl150_760])])).
fof(f39478, plain, ((e22 = op2(e21, e23)) | (~ spl150_267 | ~ spl150_347 | ~ spl150_760)), inference(forward_demodulation, [], [f29098, f6957])).
fof(f6957, plain, ((e21 = h20(e13)) | ~ spl150_760), inference(avatar_component_clause, [], [f6956])).
fof(f29098, plain, ((e22 = op2(h20(e13), e23)) | (~ spl150_267 | ~ spl150_347)), inference(forward_demodulation, [], [f29097, f3857])).
fof(f29097, plain, ((op2(e25, e24) = op2(h20(e13), e23)) | ~ spl150_347), inference(forward_demodulation, [], [f5589, f4190])).
fof(f5589, plain, (op2(e25, op2(e23, e23)) = op2(h20(e13), e23)), inference(backward_demodulation, [], [f884, f1600])).
fof(f1600, plain, (op2(e25, e23) = h20(e13)), inference(cnf_transformation, [], [f45])).
fof(f45, plain, ((op2(op2(e25, e23), op2(e25, e23)) = h20(e14)) & (op2(e25, e23) = h20(e13)) & (op2(e25, op2(op2(e25, e23), op2(e25, e23))) = h20(e12)) & (h20(e10) = op2(op2(e25, op2(op2(e25, e23), op2(e25, e23))), op2(e25, op2(op2(e25, e23), op2(e25, e23))))) & (e25 = h20(e15)) & (e23 = h20(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG027+1.p', ax45)).
fof(f884, plain, (op2(op2(e25, e23), e23) = op2(e25, op2(e23, e23))), inference(cnf_transformation, [], [f6])).
fof(f40653, plain, (~ spl150_413 | ~ spl150_725), inference(avatar_split_clause, [], [f39745, f6771, f4463])).
fof(f4463, plain, (spl150_413 <=> (e24 = op2(e21, e24))), introduced(avatar_definition, [new_symbols(naming, [spl150_413])])).
fof(f39745, plain, (~ (e24 = op2(e21, e24)) | ~ spl150_725), inference(forward_demodulation, [], [f5607, f6772])).
fof(f5607, plain, ~ (op2(e21, e24) = h21(e13)), inference(backward_demodulation, [], [f1286, f1606])).
fof(f1606, plain, (op2(e20, e24) = h21(e13)), inference(cnf_transformation, [], [f46])).
fof(f46, plain, ((op2(op2(e20, e24), op2(e20, e24)) = h21(e14)) & (op2(e20, e24) = h21(e13)) & (op2(e20, op2(op2(e20, e24), op2(e20, e24))) = h21(e12)) & (h21(e10) = op2(op2(e20, op2(op2(e20, e24), op2(e20, e24))), op2(e20, op2(op2(e20, e24), op2(e20, e24))))) & (e20 = h21(e15)) & (e24 = h21(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG027+1.p', ax46)).
fof(f1286, plain, ~ (op2(e20, e24) = op2(e21, e24)), inference(cnf_transformation, [], [f18])).
fof(f40628, plain, (~ spl150_408 | ~ spl150_1336), inference(avatar_split_clause, [], [f39830, f10449, f4442])).
fof(f4442, plain, (spl150_408 <=> (e25 = op2(e21, e25))), introduced(avatar_definition, [new_symbols(naming, [spl150_408])])).
fof(f39830, plain, (~ (e25 = op2(e21, e25)) | ~ spl150_1336), inference(forward_demodulation, [], [f5706, f10450])).
fof(f5706, plain, ~ (op2(e21, e25) = h26(e13)), inference(backward_demodulation, [], [f1301, f1636])).
fof(f1636, plain, (op2(e20, e25) = h26(e13)), inference(cnf_transformation, [], [f51])).
fof(f51, plain, ((op2(op2(e20, e25), op2(e20, e25)) = h26(e14)) & (op2(e20, e25) = h26(e13)) & (op2(e20, op2(op2(e20, e25), op2(e20, e25))) = h26(e12)) & (h26(e10) = op2(op2(e20, op2(op2(e20, e25), op2(e20, e25))), op2(e20, op2(op2(e20, e25), op2(e20, e25))))) & (e20 = h26(e15)) & (e25 = h26(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG027+1.p', ax51)).
fof(f1301, plain, ~ (op2(e20, e25) = op2(e21, e25)), inference(cnf_transformation, [], [f18])).
fof(f40561, plain, (spl150_46 | ~ spl150_49 | ~ spl150_89 | ~ spl150_196), inference(avatar_split_clause, [], [f40560, f3559, f3113, f2947, f2934])).
fof(f40560, plain, ((e13 = op1(e14, e14)) | (~ spl150_49 | ~ spl150_89 | ~ spl150_196)), inference(forward_demodulation, [], [f40555, f3561])).
fof(f40555, plain, ((op1(e10, e13) = op1(e14, e14)) | (~ spl150_49 | ~ spl150_89)), inference(backward_demodulation, [], [f24344, f2949])).
fof(f24344, plain, ((op1(e14, e14) = op1(op1(e14, e13), e13)) | ~ spl150_89), inference(forward_demodulation, [], [f565, f3115])).
fof(f565, plain, (op1(op1(e14, e13), e13) = op1(e14, op1(e13, e13))), inference(cnf_transformation, [], [f2])).
fof(f2, plain, ((op1(op1(e15, e15), e15) = op1(e15, op1(e15, e15))) & (op1(op1(e15, e15), e14) = op1(e15, op1(e15, e14))) & (op1(op1(e15, e15), e13) = op1(e15, op1(e15, e13))) & (op1(op1(e15, e15), e12) = op1(e15, op1(e15, e12))) & (op1(op1(e15, e15), e11) = op1(e15, op1(e15, e11))) & (op1(op1(e15, e15), e10) = op1(e15, op1(e15, e10))) & (op1(op1(e15, e14), e15) = op1(e15, op1(e14, e15))) & (op1(op1(e15, e14), e14) = op1(e15, op1(e14, e14))) & (op1(op1(e15, e14), e13) = op1(e15, op1(e14, e13))) & (op1(op1(e15, e14), e12) = op1(e15, op1(e14, e12))) & (op1(op1(e15, e14), e11) = op1(e15, op1(e14, e11))) & (op1(op1(e15, e14), e10) = op1(e15, op1(e14, e10))) & (op1(op1(e15, e13), e15) = op1(e15, op1(e13, e15))) & (op1(op1(e15, e13), e14) = op1(e15, op1(e13, e14))) & (op1(op1(e15, e13), e13) = op1(e15, op1(e13, e13))) & (op1(op1(e15, e13), e12) = op1(e15, op1(e13, e12))) & (op1(op1(e15, e13), e11) = op1(e15, op1(e13, e11))) & (op1(op1(e15, e13), e10) = op1(e15, op1(e13, e10))) & (op1(op1(e15, e12), e15) = op1(e15, op1(e12, e15))) & (op1(op1(e15, e12), e14) = op1(e15, op1(e12, e14))) & (op1(op1(e15, e12), e13) = op1(e15, op1(e12, e13))) & (op1(op1(e15, e12), e12) = op1(e15, op1(e12, e12))) & (op1(op1(e15, e12), e11) = op1(e15, op1(e12, e11))) & (op1(op1(e15, e12), e10) = op1(e15, op1(e12, e10))) & (op1(op1(e15, e11), e15) = op1(e15, op1(e11, e15))) & (op1(op1(e15, e11), e14) = op1(e15, op1(e11, e14))) & (op1(op1(e15, e11), e13) = op1(e15, op1(e11, e13))) & (op1(op1(e15, e11), e12) = op1(e15, op1(e11, e12))) & (op1(op1(e15, e11), e11) = op1(e15, op1(e11, e11))) & (op1(op1(e15, e11), e10) = op1(e15, op1(e11, e10))) & (op1(op1(e15, e10), e15) = op1(e15, op1(e10, e15))) & (op1(op1(e15, e10), e14) = op1(e15, op1(e10, e14))) & (op1(op1(e15, e10), e13) = op1(e15, op1(e10, e13))) & (op1(op1(e15, e10), e12) = op1(e15, op1(e10, e12))) & (op1(op1(e15, e10), e11) = op1(e15, op1(e10, e11))) & (op1(op1(e15, e10), e10) = op1(e15, op1(e10, e10))) & (op1(op1(e14, e15), e15) = op1(e14, op1(e15, e15))) & (op1(op1(e14, e15), e14) = op1(e14, op1(e15, e14))) & (op1(op1(e14, e15), e13) = op1(e14, op1(e15, e13))) & (op1(op1(e14, e15), e12) = op1(e14, op1(e15, e12))) & (op1(op1(e14, e15), e11) = op1(e14, op1(e15, e11))) & (op1(op1(e14, e15), e10) = op1(e14, op1(e15, e10))) & (op1(op1(e14, e14), e15) = op1(e14, op1(e14, e15))) & (op1(op1(e14, e14), e14) = op1(e14, op1(e14, e14))) & (op1(op1(e14, e14), e13) = op1(e14, op1(e14, e13))) & (op1(op1(e14, e14), e12) = op1(e14, op1(e14, e12))) & (op1(op1(e14, e14), e11) = op1(e14, op1(e14, e11))) & (op1(op1(e14, e14), e10) = op1(e14, op1(e14, e10))) & (op1(op1(e14, e13), e15) = op1(e14, op1(e13, e15))) & (op1(op1(e14, e13), e14) = op1(e14, op1(e13, e14))) & (op1(op1(e14, e13), e13) = op1(e14, op1(e13, e13))) & (op1(op1(e14, e13), e12) = op1(e14, op1(e13, e12))) & (op1(op1(e14, e13), e11) = op1(e14, op1(e13, e11))) & (op1(op1(e14, e13), e10) = op1(e14, op1(e13, e10))) & (op1(op1(e14, e12), e15) = op1(e14, op1(e12, e15))) & (op1(op1(e14, e12), e14) = op1(e14, op1(e12, e14))) & (op1(op1(e14, e12), e13) = op1(e14, op1(e12, e13))) & (op1(op1(e14, e12), e12) = op1(e14, op1(e12, e12))) & (op1(op1(e14, e12), e11) = op1(e14, op1(e12, e11))) & (op1(op1(e14, e12), e10) = op1(e14, op1(e12, e10))) & (op1(op1(e14, e11), e15) = op1(e14, op1(e11, e15))) & (op1(op1(e14, e11), e14) = op1(e14, op1(e11, e14))) & (op1(op1(e14, e11), e13) = op1(e14, op1(e11, e13))) & (op1(op1(e14, e11), e12) = op1(e14, op1(e11, e12))) & (op1(op1(e14, e11), e11) = op1(e14, op1(e11, e11))) & (op1(op1(e14, e11), e10) = op1(e14, op1(e11, e10))) & (op1(op1(e14, e10), e15) = op1(e14, op1(e10, e15))) & (op1(op1(e14, e10), e14) = op1(e14, op1(e10, e14))) & (op1(op1(e14, e10), e13) = op1(e14, op1(e10, e13))) & (op1(op1(e14, e10), e12) = op1(e14, op1(e10, e12))) & (op1(op1(e14, e10), e11) = op1(e14, op1(e10, e11))) & (op1(op1(e14, e10), e10) = op1(e14, op1(e10, e10))) & (op1(op1(e13, e15), e15) = op1(e13, op1(e15, e15))) & (op1(op1(e13, e15), e14) = op1(e13, op1(e15, e14))) & (op1(op1(e13, e15), e13) = op1(e13, op1(e15, e13))) & (op1(op1(e13, e15), e12) = op1(e13, op1(e15, e12))) & (op1(op1(e13, e15), e11) = op1(e13, op1(e15, e11))) & (op1(op1(e13, e15), e10) = op1(e13, op1(e15, e10))) & (op1(op1(e13, e14), e15) = op1(e13, op1(e14, e15))) & (op1(op1(e13, e14), e14) = op1(e13, op1(e14, e14))) & (op1(op1(e13, e14), e13) = op1(e13, op1(e14, e13))) & (op1(op1(e13, e14), e12) = op1(e13, op1(e14, e12))) & (op1(op1(e13, e14), e11) = op1(e13, op1(e14, e11))) & (op1(op1(e13, e14), e10) = op1(e13, op1(e14, e10))) & (op1(op1(e13, e13), e15) = op1(e13, op1(e13, e15))) & (op1(op1(e13, e13), e14) = op1(e13, op1(e13, e14))) & (op1(op1(e13, e13), e13) = op1(e13, op1(e13, e13))) & (op1(op1(e13, e13), e12) = op1(e13, op1(e13, e12))) & (op1(op1(e13, e13), e11) = op1(e13, op1(e13, e11))) & (op1(op1(e13, e13), e10) = op1(e13, op1(e13, e10))) & (op1(op1(e13, e12), e15) = op1(e13, op1(e12, e15))) & (op1(op1(e13, e12), e14) = op1(e13, op1(e12, e14))) & (op1(op1(e13, e12), e13) = op1(e13, op1(e12, e13))) & (op1(op1(e13, e12), e12) = op1(e13, op1(e12, e12))) & (op1(op1(e13, e12), e11) = op1(e13, op1(e12, e11))) & (op1(op1(e13, e12), e10) = op1(e13, op1(e12, e10))) & (op1(op1(e13, e11), e15) = op1(e13, op1(e11, e15))) & (op1(op1(e13, e11), e14) = op1(e13, op1(e11, e14))) & (op1(op1(e13, e11), e13) = op1(e13, op1(e11, e13))) & (op1(op1(e13, e11), e12) = op1(e13, op1(e11, e12))) & (op1(op1(e13, e11), e11) = op1(e13, op1(e11, e11))) & (op1(op1(e13, e11), e10) = op1(e13, op1(e11, e10))) & (op1(op1(e13, e10), e15) = op1(e13, op1(e10, e15))) & (op1(op1(e13, e10), e14) = op1(e13, op1(e10, e14))) & (op1(op1(e13, e10), e13) = op1(e13, op1(e10, e13))) & (op1(op1(e13, e10), e12) = op1(e13, op1(e10, e12))) & (op1(op1(e13, e10), e11) = op1(e13, op1(e10, e11))) & (op1(op1(e13, e10), e10) = op1(e13, op1(e10, e10))) & (op1(op1(e12, e15), e15) = op1(e12, op1(e15, e15))) & (op1(op1(e12, e15), e14) = op1(e12, op1(e15, e14))) & (op1(op1(e12, e15), e13) = op1(e12, op1(e15, e13))) & (op1(op1(e12, e15), e12) = op1(e12, op1(e15, e12))) & (op1(op1(e12, e15), e11) = op1(e12, op1(e15, e11))) & (op1(op1(e12, e15), e10) = op1(e12, op1(e15, e10))) & (op1(op1(e12, e14), e15) = op1(e12, op1(e14, e15))) & (op1(op1(e12, e14), e14) = op1(e12, op1(e14, e14))) & (op1(op1(e12, e14), e13) = op1(e12, op1(e14, e13))) & (op1(op1(e12, e14), e12) = op1(e12, op1(e14, e12))) & (op1(op1(e12, e14), e11) = op1(e12, op1(e14, e11))) & (op1(op1(e12, e14), e10) = op1(e12, op1(e14, e10))) & (op1(op1(e12, e13), e15) = op1(e12, op1(e13, e15))) & (op1(op1(e12, e13), e14) = op1(e12, op1(e13, e14))) & (op1(op1(e12, e13), e13) = op1(e12, op1(e13, e13))) & (op1(op1(e12, e13), e12) = op1(e12, op1(e13, e12))) & (op1(op1(e12, e13), e11) = op1(e12, op1(e13, e11))) & (op1(op1(e12, e13), e10) = op1(e12, op1(e13, e10))) & (op1(op1(e12, e12), e15) = op1(e12, op1(e12, e15))) & (op1(op1(e12, e12), e14) = op1(e12, op1(e12, e14))) & (op1(op1(e12, e12), e13) = op1(e12, op1(e12, e13))) & (op1(op1(e12, e12), e12) = op1(e12, op1(e12, e12))) & (op1(op1(e12, e12), e11) = op1(e12, op1(e12, e11))) & (op1(op1(e12, e12), e10) = op1(e12, op1(e12, e10))) & (op1(op1(e12, e11), e15) = op1(e12, op1(e11, e15))) & (op1(op1(e12, e11), e14) = op1(e12, op1(e11, e14))) & (op1(op1(e12, e11), e13) = op1(e12, op1(e11, e13))) & (op1(op1(e12, e11), e12) = op1(e12, op1(e11, e12))) & (op1(op1(e12, e11), e11) = op1(e12, op1(e11, e11))) & (op1(op1(e12, e11), e10) = op1(e12, op1(e11, e10))) & (op1(op1(e12, e10), e15) = op1(e12, op1(e10, e15))) & (op1(op1(e12, e10), e14) = op1(e12, op1(e10, e14))) & (op1(op1(e12, e10), e13) = op1(e12, op1(e10, e13))) & (op1(op1(e12, e10), e12) = op1(e12, op1(e10, e12))) & (op1(op1(e12, e10), e11) = op1(e12, op1(e10, e11))) & (op1(op1(e12, e10), e10) = op1(e12, op1(e10, e10))) & (op1(op1(e11, e15), e15) = op1(e11, op1(e15, e15))) & (op1(op1(e11, e15), e14) = op1(e11, op1(e15, e14))) & (op1(op1(e11, e15), e13) = op1(e11, op1(e15, e13))) & (op1(op1(e11, e15), e12) = op1(e11, op1(e15, e12))) & (op1(op1(e11, e15), e11) = op1(e11, op1(e15, e11))) & (op1(op1(e11, e15), e10) = op1(e11, op1(e15, e10))) & (op1(op1(e11, e14), e15) = op1(e11, op1(e14, e15))) & (op1(op1(e11, e14), e14) = op1(e11, op1(e14, e14))) & (op1(op1(e11, e14), e13) = op1(e11, op1(e14, e13))) & (op1(op1(e11, e14), e12) = op1(e11, op1(e14, e12))) & (op1(op1(e11, e14), e11) = op1(e11, op1(e14, e11))) & (op1(op1(e11, e14), e10) = op1(e11, op1(e14, e10))) & (op1(op1(e11, e13), e15) = op1(e11, op1(e13, e15))) & (op1(op1(e11, e13), e14) = op1(e11, op1(e13, e14))) & (op1(op1(e11, e13), e13) = op1(e11, op1(e13, e13))) & (op1(op1(e11, e13), e12) = op1(e11, op1(e13, e12))) & (op1(op1(e11, e13), e11) = op1(e11, op1(e13, e11))) & (op1(op1(e11, e13), e10) = op1(e11, op1(e13, e10))) & (op1(op1(e11, e12), e15) = op1(e11, op1(e12, e15))) & (op1(op1(e11, e12), e14) = op1(e11, op1(e12, e14))) & (op1(op1(e11, e12), e13) = op1(e11, op1(e12, e13))) & (op1(op1(e11, e12), e12) = op1(e11, op1(e12, e12))) & (op1(op1(e11, e12), e11) = op1(e11, op1(e12, e11))) & (op1(op1(e11, e12), e10) = op1(e11, op1(e12, e10))) & (op1(op1(e11, e11), e15) = op1(e11, op1(e11, e15))) & (op1(op1(e11, e11), e14) = op1(e11, op1(e11, e14))) & (op1(op1(e11, e11), e13) = op1(e11, op1(e11, e13))) & (op1(op1(e11, e11), e12) = op1(e11, op1(e11, e12))) & (op1(op1(e11, e11), e11) = op1(e11, op1(e11, e11))) & (op1(op1(e11, e11), e10) = op1(e11, op1(e11, e10))) & (op1(op1(e11, e10), e15) = op1(e11, op1(e10, e15))) & (op1(op1(e11, e10), e14) = op1(e11, op1(e10, e14))) & (op1(op1(e11, e10), e13) = op1(e11, op1(e10, e13))) & (op1(op1(e11, e10), e12) = op1(e11, op1(e10, e12))) & (op1(op1(e11, e10), e11) = op1(e11, op1(e10, e11))) & (op1(op1(e11, e10), e10) = op1(e11, op1(e10, e10))) & (op1(op1(e10, e15), e15) = op1(e10, op1(e15, e15))) & (op1(op1(e10, e15), e14) = op1(e10, op1(e15, e14))) & (op1(op1(e10, e15), e13) = op1(e10, op1(e15, e13))) & (op1(op1(e10, e15), e12) = op1(e10, op1(e15, e12))) & (op1(op1(e10, e15), e11) = op1(e10, op1(e15, e11))) & (op1(op1(e10, e15), e10) = op1(e10, op1(e15, e10))) & (op1(op1(e10, e14), e15) = op1(e10, op1(e14, e15))) & (op1(op1(e10, e14), e14) = op1(e10, op1(e14, e14))) & (op1(op1(e10, e14), e13) = op1(e10, op1(e14, e13))) & (op1(op1(e10, e14), e12) = op1(e10, op1(e14, e12))) & (op1(op1(e10, e14), e11) = op1(e10, op1(e14, e11))) & (op1(op1(e10, e14), e10) = op1(e10, op1(e14, e10))) & (op1(op1(e10, e13), e15) = op1(e10, op1(e13, e15))) & (op1(op1(e10, e13), e14) = op1(e10, op1(e13, e14))) & (op1(op1(e10, e13), e13) = op1(e10, op1(e13, e13))) & (op1(op1(e10, e13), e12) = op1(e10, op1(e13, e12))) & (op1(op1(e10, e13), e11) = op1(e10, op1(e13, e11))) & (op1(op1(e10, e13), e10) = op1(e10, op1(e13, e10))) & (op1(op1(e10, e12), e15) = op1(e10, op1(e12, e15))) & (op1(op1(e10, e12), e14) = op1(e10, op1(e12, e14))) & (op1(op1(e10, e12), e13) = op1(e10, op1(e12, e13))) & (op1(op1(e10, e12), e12) = op1(e10, op1(e12, e12))) & (op1(op1(e10, e12), e11) = op1(e10, op1(e12, e11))) & (op1(op1(e10, e12), e10) = op1(e10, op1(e12, e10))) & (op1(op1(e10, e11), e15) = op1(e10, op1(e11, e15))) & (op1(op1(e10, e11), e14) = op1(e10, op1(e11, e14))) & (op1(op1(e10, e11), e13) = op1(e10, op1(e11, e13))) & (op1(op1(e10, e11), e12) = op1(e10, op1(e11, e12))) & (op1(op1(e10, e11), e11) = op1(e10, op1(e11, e11))) & (op1(op1(e10, e11), e10) = op1(e10, op1(e11, e10))) & (op1(op1(e10, e10), e15) = op1(e10, op1(e10, e15))) & (op1(op1(e10, e10), e14) = op1(e10, op1(e10, e14))) & (op1(op1(e10, e10), e13) = op1(e10, op1(e10, e13))) & (op1(op1(e10, e10), e12) = op1(e10, op1(e10, e12))) & (op1(op1(e10, e10), e11) = op1(e10, op1(e10, e11))) & (op1(op1(e10, e10), e10) = op1(e10, op1(e10, e10)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG027+1.p', ax2)).
fof(f40520, plain, (spl150_60 | ~ spl150_23 | ~ spl150_149 | ~ spl150_156), inference(avatar_split_clause, [], [f40515, f3392, f3363, f2838, f2992])).
fof(f40515, plain, ((e15 = op1(e14, e12)) | (~ spl150_23 | ~ spl150_149 | ~ spl150_156)), inference(backward_demodulation, [], [f39285, f3365])).
fof(f39285, plain, ((e15 = op1(op1(e11, e15), e12)) | (~ spl150_23 | ~ spl150_156)), inference(backward_demodulation, [], [f39266, f3394])).
fof(f39266, plain, ((op1(e11, e14) = op1(op1(e11, e15), e12)) | ~ spl150_23), inference(forward_demodulation, [], [f468, f2840])).
fof(f468, plain, (op1(op1(e11, e15), e12) = op1(e11, op1(e15, e12))), inference(cnf_transformation, [], [f2])).
fof(f40473, plain, (spl150_133 | ~ spl150_217 | ~ spl150_249), inference(avatar_contradiction_clause, [], [f40472])).
fof(f40472, plain, ($false | (spl150_133 | ~ spl150_217 | ~ spl150_249)), inference(subsumption_resolution, [], [f40468, f3298])).
fof(f3298, plain, (~ (e10 = op1(e12, e11)) | spl150_133), inference(avatar_component_clause, [], [f3297])).
fof(f3297, plain, (spl150_133 <=> (e10 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl150_133])])).
fof(f40468, plain, ((e10 = op1(e12, e11)) | (~ spl150_217 | ~ spl150_249)), inference(backward_demodulation, [], [f37273, f3782])).
fof(f3782, plain, ((e12 = inv1(e11)) | ~ spl150_249), inference(avatar_component_clause, [], [f3780])).
fof(f3780, plain, (spl150_249 <=> (e12 = inv1(e11))), introduced(avatar_definition, [new_symbols(naming, [spl150_249])])).
fof(f40427, plain, (spl150_307 | ~ spl150_475 | ~ spl150_497), inference(avatar_split_clause, [], [f40424, f4813, f4722, f4022])).
fof(f4722, plain, (spl150_475 <=> (e20 = unit2)), introduced(avatar_definition, [new_symbols(naming, [spl150_475])])).
fof(f4813, plain, (spl150_497 <=> (e24 = inv2(e23))), introduced(avatar_definition, [new_symbols(naming, [spl150_497])])).
fof(f40424, plain, ((e20 = op2(e24, e23)) | (~ spl150_475 | ~ spl150_497)), inference(backward_demodulation, [], [f39681, f4815])).
fof(f4815, plain, ((e24 = inv2(e23)) | ~ spl150_497), inference(avatar_component_clause, [], [f4813])).
fof(f39681, plain, ((e20 = op2(inv2(e23), e23)) | ~ spl150_475), inference(forward_demodulation, [], [f919, f4724])).
fof(f4724, plain, ((e20 = unit2) | ~ spl150_475), inference(avatar_component_clause, [], [f4722])).
fof(f919, plain, (unit2 = op2(inv2(e23), e23)), inference(cnf_transformation, [], [f8])).
fof(f8, plain, (((e25 = inv2(e25)) | (e24 = inv2(e25)) | (e23 = inv2(e25)) | (e22 = inv2(e25)) | (e21 = inv2(e25)) | (e20 = inv2(e25))) & ((e25 = inv2(e24)) | (e24 = inv2(e24)) | (e23 = inv2(e24)) | (e22 = inv2(e24)) | (e21 = inv2(e24)) | (e20 = inv2(e24))) & ((e25 = inv2(e23)) | (e24 = inv2(e23)) | (e23 = inv2(e23)) | (e22 = inv2(e23)) | (e21 = inv2(e23)) | (e20 = inv2(e23))) & ((e25 = inv2(e22)) | (e24 = inv2(e22)) | (e23 = inv2(e22)) | (e22 = inv2(e22)) | (e21 = inv2(e22)) | (e20 = inv2(e22))) & ((e25 = inv2(e21)) | (e24 = inv2(e21)) | (e23 = inv2(e21)) | (e22 = inv2(e21)) | (e21 = inv2(e21)) | (e20 = inv2(e21))) & ((e25 = inv2(e20)) | (e24 = inv2(e20)) | (e23 = inv2(e20)) | (e22 = inv2(e20)) | (e21 = inv2(e20)) | (e20 = inv2(e20))) & (unit2 = op2(inv2(e25), e25)) & (unit2 = op2(e25, inv2(e25))) & (unit2 = op2(inv2(e24), e24)) & (unit2 = op2(e24, inv2(e24))) & (unit2 = op2(inv2(e23), e23)) & (unit2 = op2(e23, inv2(e23))) & (unit2 = op2(inv2(e22), e22)) & (unit2 = op2(e22, inv2(e22))) & (unit2 = op2(inv2(e21), e21)) & (unit2 = op2(e21, inv2(e21))) & (unit2 = op2(inv2(e20), e20)) & (unit2 = op2(e20, inv2(e20)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG027+1.p', ax8)).
fof(f40411, plain, (spl150_427 | ~ spl150_475 | ~ spl150_506), inference(avatar_split_clause, [], [f40408, f4851, f4722, f4522])).
fof(f4851, plain, (spl150_506 <=> (e21 = inv2(e21))), introduced(avatar_definition, [new_symbols(naming, [spl150_506])])).
fof(f40408, plain, ((e20 = op2(e21, e21)) | (~ spl150_475 | ~ spl150_506)), inference(backward_demodulation, [], [f39769, f4853])).
fof(f4853, plain, ((e21 = inv2(e21)) | ~ spl150_506), inference(avatar_component_clause, [], [f4851])).
fof(f39769, plain, ((e20 = op2(e21, inv2(e21))) | ~ spl150_475), inference(forward_demodulation, [], [f914, f4724])).
fof(f914, plain, (unit2 = op2(e21, inv2(e21))), inference(cnf_transformation, [], [f8])).
fof(f40398, plain, (~ spl150_505 | ~ spl150_511), inference(avatar_split_clause, [], [f40387, f4872, f4847])).
fof(f4847, plain, (spl150_505 <=> (e20 = inv2(e21))), introduced(avatar_definition, [new_symbols(naming, [spl150_505])])).
fof(f4872, plain, (spl150_511 <=> (e20 = inv2(e20))), introduced(avatar_definition, [new_symbols(naming, [spl150_511])])).
fof(f40387, plain, (~ (e20 = inv2(e21)) | ~ spl150_511), inference(backward_demodulation, [], [f1031, f4874])).
fof(f4874, plain, ((e20 = inv2(e20)) | ~ spl150_511), inference(avatar_component_clause, [], [f4872])).
fof(f1031, plain, ~ (inv2(e20) = inv2(e21)), inference(cnf_transformation, [], [f16])).
fof(f16, plain, (~ (inv2(e24) = inv2(e25)) & ~ (inv2(e23) = inv2(e25)) & ~ (inv2(e23) = inv2(e24)) & ~ (inv2(e22) = inv2(e25)) & ~ (inv2(e22) = inv2(e24)) & ~ (inv2(e22) = inv2(e23)) & ~ (inv2(e21) = inv2(e25)) & ~ (inv2(e21) = inv2(e24)) & ~ (inv2(e21) = inv2(e23)) & ~ (inv2(e21) = inv2(e22)) & ~ (inv2(e20) = inv2(e25)) & ~ (inv2(e20) = inv2(e24)) & ~ (inv2(e20) = inv2(e23)) & ~ (inv2(e20) = inv2(e22)) & ~ (inv2(e20) = inv2(e21))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG027+1.p', ax16)).
fof(f40273, plain, (~ spl150_211 | ~ spl150_469 | spl150_1973), inference(avatar_contradiction_clause, [], [f40272])).
fof(f40272, plain, ($false | (~ spl150_211 | ~ spl150_469 | spl150_1973)), inference(subsumption_resolution, [], [f40271, f4699])).
fof(f4699, plain, ((e20 = op2(e20, e20)) | ~ spl150_469), inference(avatar_component_clause, [], [f4697])).
fof(f4697, plain, (spl150_469 <=> (e20 = op2(e20, e20))), introduced(avatar_definition, [new_symbols(naming, [spl150_469])])).
fof(f40271, plain, (~ (e20 = op2(e20, e20)) | (~ spl150_211 | spl150_1973)), inference(forward_demodulation, [], [f40270, f5365])).
fof(f40270, plain, (~ (op2(e20, e20) = h10(e10)) | (~ spl150_211 | spl150_1973)), inference(forward_demodulation, [], [f15271, f3624])).
fof(f3624, plain, ((e10 = op1(e10, e10)) | ~ spl150_211), inference(avatar_component_clause, [], [f3622])).
fof(f3622, plain, (spl150_211 <=> (e10 = op1(e10, e10))), introduced(avatar_definition, [new_symbols(naming, [spl150_211])])).
fof(f15271, plain, (~ (op2(e20, e20) = h10(op1(e10, e10))) | spl150_1973), inference(avatar_component_clause, [], [f15269])).
fof(f15269, plain, (spl150_1973 <=> (op2(e20, e20) = h10(op1(e10, e10)))), introduced(avatar_definition, [new_symbols(naming, [spl150_1973])])).
fof(f40044, plain, (~ spl150_272 | ~ spl150_347 | ~ spl150_710 | spl150_780 | ~ spl150_954), inference(avatar_contradiction_clause, [], [f40043])).
fof(f40043, plain, ($false | (~ spl150_272 | ~ spl150_347 | ~ spl150_710 | spl150_780 | ~ spl150_954)), inference(subsumption_resolution, [], [f40042, f7064])).
fof(f7064, plain, (~ (e21 = h19(e13)) | spl150_780), inference(avatar_component_clause, [], [f7062])).
fof(f7062, plain, (spl150_780 <=> (e21 = h19(e13))), introduced(avatar_definition, [new_symbols(naming, [spl150_780])])).
fof(f40042, plain, ((e21 = h19(e13)) | (~ spl150_272 | ~ spl150_347 | ~ spl150_710 | ~ spl150_954)), inference(forward_demodulation, [], [f40041, f3878])).
fof(f40041, plain, ((op2(e25, e23) = h19(e13)) | (~ spl150_347 | ~ spl150_710 | ~ spl150_954)), inference(forward_demodulation, [], [f34934, f6693])).
fof(f6693, plain, ((e23 = h22(e13)) | ~ spl150_710), inference(avatar_component_clause, [], [f6692])).
fof(f6692, plain, (spl150_710 <=> (e23 = h22(e13))), introduced(avatar_definition, [new_symbols(naming, [spl150_710])])).
fof(f34934, plain, ((h19(e13) = op2(e25, h22(e13))) | (~ spl150_347 | ~ spl150_954)), inference(backward_demodulation, [], [f20099, f34923])).
fof(f34923, plain, ((h19(e13) = h24(e13)) | ~ spl150_347), inference(backward_demodulation, [], [f1624, f26724])).
fof(f26724, plain, ((op2(e23, e24) = h19(e13)) | ~ spl150_347), inference(forward_demodulation, [], [f26723, f1594])).
fof(f26723, plain, ((op2(e23, e24) = op2(e24, e23)) | ~ spl150_347), inference(forward_demodulation, [], [f812, f4190])).
fof(f812, plain, (op2(op2(e23, e23), e23) = op2(e23, op2(e23, e23))), inference(cnf_transformation, [], [f6])).
fof(f1624, plain, (op2(e23, e24) = h24(e13)), inference(cnf_transformation, [], [f49])).
fof(f49, plain, ((op2(op2(e23, e24), op2(e23, e24)) = h24(e14)) & (op2(e23, e24) = h24(e13)) & (op2(e23, op2(op2(e23, e24), op2(e23, e24))) = h24(e12)) & (h24(e10) = op2(op2(e23, op2(op2(e23, e24), op2(e23, e24))), op2(e23, op2(op2(e23, e24), op2(e23, e24))))) & (e23 = h24(e15)) & (e24 = h24(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG027+1.p', ax49)).
fof(f20099, plain, ((h24(e13) = op2(e25, h22(e13))) | ~ spl150_954), inference(forward_demodulation, [], [f20080, f1624])).
fof(f20080, plain, ((op2(e23, e24) = op2(e25, h22(e13))) | ~ spl150_954), inference(backward_demodulation, [], [f5643, f7986])).
fof(f5643, plain, (op2(h10(e13), e24) = op2(e25, h22(e13))), inference(backward_demodulation, [], [f5328, f1612])).
fof(f1612, plain, (op2(e21, e24) = h22(e13)), inference(cnf_transformation, [], [f47])).
fof(f47, plain, ((op2(op2(e21, e24), op2(e21, e24)) = h22(e14)) & (op2(e21, e24) = h22(e13)) & (op2(e21, op2(op2(e21, e24), op2(e21, e24))) = h22(e12)) & (h22(e10) = op2(op2(e21, op2(op2(e21, e24), op2(e21, e24))), op2(e21, op2(op2(e21, e24), op2(e21, e24))))) & (e21 = h22(e15)) & (e24 = h22(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG027+1.p', ax47)).
fof(f5328, plain, (op2(e25, op2(e21, e24)) = op2(h10(e13), e24)), inference(backward_demodulation, [], [f873, f1540])).
fof(f873, plain, (op2(op2(e25, e21), e24) = op2(e25, op2(e21, e24))), inference(cnf_transformation, [], [f6])).
fof(f39770, plain, (spl150_511 | ~ spl150_475), inference(avatar_split_clause, [], [f38344, f4722, f4872])).
fof(f38344, plain, ((e20 = inv2(e20)) | ~ spl150_475), inference(backward_demodulation, [], [f988, f4724])).
fof(f988, plain, (unit2 = inv2(unit2)), inference(cnf_transformation, [], [f13])).
fof(f13, plain, (unit2 = inv2(unit2)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG027+1.p', ax13)).
fof(f39202, plain, (spl150_116 | ~ spl150_127 | ~ spl150_137 | ~ spl150_206), inference(avatar_split_clause, [], [f38818, f3601, f3313, f3272, f3226])).
fof(f38818, plain, ((e11 = op1(e12, e14)) | (~ spl150_127 | ~ spl150_137 | ~ spl150_206)), inference(backward_demodulation, [], [f34048, f3315])).
fof(f34048, plain, ((e11 = op1(e12, op1(e12, e11))) | (~ spl150_127 | ~ spl150_206)), inference(backward_demodulation, [], [f26170, f3603])).
fof(f26170, plain, ((op1(e10, e11) = op1(e12, op1(e12, e11))) | ~ spl150_127), inference(forward_demodulation, [], [f485, f3274])).
fof(f485, plain, (op1(op1(e12, e12), e11) = op1(e12, op1(e12, e11))), inference(cnf_transformation, [], [f2])).
fof(f39074, plain, (spl150_49 | ~ spl150_79 | ~ spl150_526), inference(avatar_split_clause, [], [f39073, f4993, f3072, f2947])).
fof(f4993, plain, (spl150_526 <=> (op1(e13, e14) = op1(e14, e13))), introduced(avatar_definition, [new_symbols(naming, [spl150_526])])).
fof(f39073, plain, ((e10 = op1(e14, e13)) | (~ spl150_79 | ~ spl150_526)), inference(forward_demodulation, [], [f4994, f3074])).
fof(f4994, plain, ((op1(e13, e14) = op1(e14, e13)) | ~ spl150_526), inference(avatar_component_clause, [], [f4993])).
fof(f39053, plain, (spl150_149 | ~ spl150_102 | ~ spl150_137 | ~ spl150_159), inference(avatar_split_clause, [], [f39052, f3405, f3313, f3167, f3363])).
fof(f39052, plain, ((e14 = op1(e11, e15)) | (~ spl150_102 | ~ spl150_137 | ~ spl150_159)), inference(forward_demodulation, [], [f39051, f3315])).
fof(f39051, plain, ((op1(e11, e15) = op1(e12, e11)) | (~ spl150_102 | ~ spl150_159)), inference(forward_demodulation, [], [f39050, f3407])).
fof(f39050, plain, ((op1(e11, e15) = op1(op1(e11, e13), e11)) | ~ spl150_102), inference(forward_demodulation, [], [f455, f3169])).
fof(f455, plain, (op1(op1(e11, e13), e11) = op1(e11, op1(e13, e11))), inference(cnf_transformation, [], [f2])).
fof(f38758, plain, (spl150_75 | ~ spl150_112 | ~ spl150_159 | ~ spl150_166), inference(avatar_split_clause, [], [f38746, f3434, f3405, f3209, f3055])).
fof(f38746, plain, ((e12 = op1(e13, e15)) | (~ spl150_112 | ~ spl150_159 | ~ spl150_166)), inference(backward_demodulation, [], [f37008, f3436])).
fof(f37008, plain, ((e12 = op1(op1(e11, e12), e15)) | (~ spl150_112 | ~ spl150_159)), inference(forward_demodulation, [], [f37007, f3407])).
fof(f37007, plain, ((op1(e11, e13) = op1(op1(e11, e12), e15)) | ~ spl150_112), inference(forward_demodulation, [], [f453, f3211])).
fof(f453, plain, (op1(op1(e11, e12), e15) = op1(e11, op1(e12, e15))), inference(cnf_transformation, [], [f2])).
fof(f38757, plain, (spl150_92 | ~ spl150_127 | ~ spl150_166 | ~ spl150_176), inference(avatar_split_clause, [], [f38743, f3476, f3434, f3272, f3126])).
fof(f38743, plain, ((e11 = op1(e13, e12)) | (~ spl150_127 | ~ spl150_166 | ~ spl150_176)), inference(backward_demodulation, [], [f32630, f3436])).
fof(f32630, plain, ((e11 = op1(op1(e11, e12), e12)) | (~ spl150_127 | ~ spl150_176)), inference(forward_demodulation, [], [f26143, f3478])).
fof(f26143, plain, ((op1(e11, e10) = op1(op1(e11, e12), e12)) | ~ spl150_127), inference(forward_demodulation, [], [f450, f3274])).
fof(f450, plain, (op1(op1(e11, e12), e12) = op1(e11, op1(e12, e12))), inference(cnf_transformation, [], [f2])).
fof(f38728, plain, (spl150_166 | ~ spl150_159 | ~ spl150_169 | ~ spl150_196), inference(avatar_split_clause, [], [f38727, f3559, f3447, f3405, f3434])).
fof(f38727, plain, ((e13 = op1(e11, e12)) | (~ spl150_159 | ~ spl150_169 | ~ spl150_196)), inference(forward_demodulation, [], [f38713, f3561])).
fof(f38713, plain, ((op1(e10, e13) = op1(e11, e12)) | (~ spl150_159 | ~ spl150_169)), inference(backward_demodulation, [], [f37204, f3449])).
fof(f37204, plain, ((op1(e11, e12) = op1(op1(e11, e11), e13)) | ~ spl150_159), inference(forward_demodulation, [], [f445, f3407])).
fof(f445, plain, (op1(op1(e11, e11), e13) = op1(e11, op1(e11, e13))), inference(cnf_transformation, [], [f2])).
fof(f38646, plain, (spl150_97 | ~ spl150_217 | ~ spl150_250), inference(avatar_contradiction_clause, [], [f38645])).
fof(f38645, plain, ($false | (spl150_97 | ~ spl150_217 | ~ spl150_250)), inference(subsumption_resolution, [], [f38641, f3148])).
fof(f3148, plain, (~ (e10 = op1(e13, e11)) | spl150_97), inference(avatar_component_clause, [], [f3147])).
fof(f3147, plain, (spl150_97 <=> (e10 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl150_97])])).
fof(f38641, plain, ((e10 = op1(e13, e11)) | (~ spl150_217 | ~ spl150_250)), inference(backward_demodulation, [], [f37273, f3786])).
fof(f3786, plain, ((e13 = inv1(e11)) | ~ spl150_250), inference(avatar_component_clause, [], [f3784])).
fof(f3784, plain, (spl150_250 <=> (e13 = inv1(e11))), introduced(avatar_definition, [new_symbols(naming, [spl150_250])])).
fof(f38616, plain, (spl150_319 | ~ spl150_329 | ~ spl150_475 | ~ spl150_509 | ~ spl150_1039), inference(avatar_split_clause, [], [f38518, f8447, f4863, f4722, f4113, f4072])).
fof(f4072, plain, (spl150_319 <=> (e20 = op2(e24, e21))), introduced(avatar_definition, [new_symbols(naming, [spl150_319])])).
fof(f4863, plain, (spl150_509 <=> (e24 = inv2(e21))), introduced(avatar_definition, [new_symbols(naming, [spl150_509])])).
fof(f38518, plain, ((e20 = op2(e24, e21)) | (~ spl150_329 | ~ spl150_475 | ~ spl150_509 | ~ spl150_1039)), inference(backward_demodulation, [], [f38393, f38500])).
fof(f38393, plain, ((e20 = op2(h4(e13), e21)) | (~ spl150_475 | ~ spl150_509 | ~ spl150_1039)), inference(backward_demodulation, [], [f38081, f4724])).
fof(f38081, plain, ((unit2 = op2(h4(e13), e21)) | (~ spl150_509 | ~ spl150_1039)), inference(forward_demodulation, [], [f38044, f35999])).
fof(f35999, plain, ((op2(e24, e21) = unit2) | ~ spl150_509), inference(backward_demodulation, [], [f915, f4865])).
fof(f4865, plain, ((e24 = inv2(e21)) | ~ spl150_509), inference(avatar_component_clause, [], [f4863])).
fof(f915, plain, (unit2 = op2(inv2(e21), e21)), inference(cnf_transformation, [], [f8])).
fof(f38044, plain, ((op2(e24, e21) = op2(h4(e13), e21)) | ~ spl150_1039), inference(backward_demodulation, [], [f5237, f8448])).
fof(f5237, plain, (op2(h4(e13), e21) = op2(e24, h6(e13))), inference(backward_demodulation, [], [f5172, f1516])).
fof(f1516, plain, (op2(e20, e21) = h6(e13)), inference(cnf_transformation, [], [f31])).
fof(f31, plain, ((op2(op2(e20, e21), op2(e20, e21)) = h6(e14)) & (op2(e20, e21) = h6(e13)) & (op2(e20, op2(op2(e20, e21), op2(e20, e21))) = h6(e12)) & (h6(e10) = op2(op2(e20, op2(op2(e20, e21), op2(e20, e21))), op2(e20, op2(op2(e20, e21), op2(e20, e21))))) & (e20 = h6(e15)) & (e21 = h6(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG027+1.p', ax31)).
fof(f5172, plain, (op2(e24, op2(e20, e21)) = op2(h4(e13), e21)), inference(backward_demodulation, [], [f828, f1504])).
fof(f828, plain, (op2(op2(e24, e20), e21) = op2(e24, op2(e20, e21))), inference(cnf_transformation, [], [f6])).
fof(f38404, plain, (spl150_434 | ~ spl150_475), inference(avatar_split_clause, [], [f38341, f4722, f4551])).
fof(f38341, plain, ((e21 = op2(e21, e20)) | ~ spl150_475), inference(backward_demodulation, [], [f902, f4724])).
fof(f902, plain, (e21 = op2(e21, unit2)), inference(cnf_transformation, [], [f7])).
fof(f7, plain, (((e25 = unit2) | (e24 = unit2) | (e23 = unit2) | (e22 = unit2) | (e21 = unit2) | (e20 = unit2)) & (e25 = op2(e25, unit2)) & (e25 = op2(unit2, e25)) & (e24 = op2(e24, unit2)) & (e24 = op2(unit2, e24)) & (e23 = op2(e23, unit2)) & (e23 = op2(unit2, e23)) & (e22 = op2(e22, unit2)) & (e22 = op2(unit2, e22)) & (e21 = op2(e21, unit2)) & (e21 = op2(unit2, e21)) & (e20 = op2(e20, unit2)) & (e20 = op2(unit2, e20))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG027+1.p', ax7)).
fof(f38403, plain, (spl150_464 | ~ spl150_475), inference(avatar_split_clause, [], [f38340, f4722, f4676])).
fof(f4676, plain, (spl150_464 <=> (e21 = op2(e20, e21))), introduced(avatar_definition, [new_symbols(naming, [spl150_464])])).
fof(f38340, plain, ((e21 = op2(e20, e21)) | ~ spl150_475), inference(backward_demodulation, [], [f901, f4724])).
fof(f901, plain, (e21 = op2(unit2, e21)), inference(cnf_transformation, [], [f7])).
fof(f38402, plain, (spl150_469 | ~ spl150_475), inference(avatar_split_clause, [], [f38339, f4722, f4697])).
fof(f38339, plain, ((e20 = op2(e20, e20)) | ~ spl150_475), inference(backward_demodulation, [], [f900, f4724])).
fof(f900, plain, (e20 = op2(e20, unit2)), inference(cnf_transformation, [], [f7])).
fof(f38165, plain, (spl150_301 | ~ spl150_385 | ~ spl150_554 | ~ spl150_656 | ~ spl150_850), inference(avatar_contradiction_clause, [], [f38164])).
fof(f38164, plain, ($false | (spl150_301 | ~ spl150_385 | ~ spl150_554 | ~ spl150_656 | ~ spl150_850)), inference(subsumption_resolution, [], [f38138, f3998])).
fof(f3998, plain, (~ (e20 = op2(e24, e24)) | spl150_301), inference(avatar_component_clause, [], [f3997])).
fof(f38138, plain, ((e20 = op2(e24, e24)) | (~ spl150_385 | ~ spl150_554 | ~ spl150_656 | ~ spl150_850)), inference(backward_demodulation, [], [f37833, f7433])).
fof(f7433, plain, ((e24 = h15(e13)) | ~ spl150_850), inference(avatar_component_clause, [], [f7432])).
fof(f7432, plain, (spl150_850 <=> (e24 = h15(e13))), introduced(avatar_definition, [new_symbols(naming, [spl150_850])])).
fof(f37833, plain, ((e20 = op2(h15(e13), e24)) | (~ spl150_385 | ~ spl150_554 | ~ spl150_656)), inference(backward_demodulation, [], [f29184, f37821])).
fof(f37821, plain, ((h15(e13) = h28(e13)) | (~ spl150_554 | ~ spl150_656)), inference(backward_demodulation, [], [f37804, f1570])).
fof(f1570, plain, (op2(e25, e22) = h15(e13)), inference(cnf_transformation, [], [f40])).
fof(f40, plain, ((op2(op2(e25, e22), op2(e25, e22)) = h15(e14)) & (op2(e25, e22) = h15(e13)) & (op2(e25, op2(op2(e25, e22), op2(e25, e22))) = h15(e12)) & (h15(e10) = op2(op2(e25, op2(op2(e25, e22), op2(e25, e22))), op2(e25, op2(op2(e25, e22), op2(e25, e22))))) & (e25 = h15(e15)) & (e22 = h15(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG027+1.p', ax40)).
fof(f37804, plain, ((op2(e25, e22) = h28(e13)) | (~ spl150_554 | ~ spl150_656)), inference(forward_demodulation, [], [f20144, f5853])).
fof(f5853, plain, ((e22 = h30(e13)) | ~ spl150_554), inference(avatar_component_clause, [], [f5852])).
fof(f5852, plain, (spl150_554 <=> (e22 = h30(e13))), introduced(avatar_definition, [new_symbols(naming, [spl150_554])])).
fof(f20144, plain, ((h28(e13) = op2(e25, h30(e13))) | ~ spl150_656), inference(forward_demodulation, [], [f20125, f1648])).
fof(f1648, plain, (op2(e22, e25) = h28(e13)), inference(cnf_transformation, [], [f53])).
fof(f53, plain, ((op2(op2(e22, e25), op2(e22, e25)) = h28(e14)) & (op2(e22, e25) = h28(e13)) & (op2(e22, op2(op2(e22, e25), op2(e22, e25))) = h28(e12)) & (h28(e10) = op2(op2(e22, op2(op2(e22, e25), op2(e22, e25))), op2(e22, op2(op2(e22, e25), op2(e22, e25))))) & (e22 = h28(e15)) & (e25 = h28(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG027+1.p', ax53)).
fof(f20125, plain, ((op2(e22, e25) = op2(e25, h30(e13))) | ~ spl150_656), inference(backward_demodulation, [], [f5802, f6399])).
fof(f5802, plain, (op2(h10(e12), e25) = op2(e25, h30(e13))), inference(backward_demodulation, [], [f5348, f1660])).
fof(f1660, plain, (op2(e24, e25) = h30(e13)), inference(cnf_transformation, [], [f55])).
fof(f55, plain, ((op2(op2(e24, e25), op2(e24, e25)) = h30(e14)) & (op2(e24, e25) = h30(e13)) & (op2(e24, op2(op2(e24, e25), op2(e24, e25))) = h30(e12)) & (h30(e10) = op2(op2(e24, op2(op2(e24, e25), op2(e24, e25))), op2(e24, op2(op2(e24, e25), op2(e24, e25))))) & (e24 = h30(e15)) & (e25 = h30(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG027+1.p', ax55)).
fof(f5348, plain, (op2(e25, op2(e24, e25)) = op2(h10(e12), e25)), inference(backward_demodulation, [], [f892, f5339])).
fof(f892, plain, (op2(op2(e25, e24), e25) = op2(e25, op2(e24, e25))), inference(cnf_transformation, [], [f6])).
fof(f29184, plain, ((e20 = op2(h28(e13), e24)) | (~ spl150_385 | ~ spl150_656)), inference(forward_demodulation, [], [f20122, f4349])).
fof(f4349, plain, ((e20 = op2(e22, e22)) | ~ spl150_385), inference(avatar_component_clause, [], [f4347])).
fof(f4347, plain, (spl150_385 <=> (e20 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl150_385])])).
fof(f20122, plain, ((op2(e22, e22) = op2(h28(e13), e24)) | ~ spl150_656), inference(backward_demodulation, [], [f5760, f6399])).
fof(f5760, plain, (op2(e22, h10(e12)) = op2(h28(e13), e24)), inference(backward_demodulation, [], [f5342, f1648])).
fof(f5342, plain, (op2(op2(e22, e25), e24) = op2(e22, h10(e12))), inference(backward_demodulation, [], [f789, f5339])).
fof(f789, plain, (op2(op2(e22, e25), e24) = op2(e22, op2(e25, e24))), inference(cnf_transformation, [], [f6])).
fof(f37742, plain, (spl150_840 | ~ spl150_385 | ~ spl150_495), inference(avatar_split_clause, [], [f37741, f4805, f4347, f7379])).
fof(f7379, plain, (spl150_840 <=> (e22 = h16(e13))), introduced(avatar_definition, [new_symbols(naming, [spl150_840])])).
fof(f4805, plain, (spl150_495 <=> (e22 = inv2(e23))), introduced(avatar_definition, [new_symbols(naming, [spl150_495])])).
fof(f37741, plain, ((e22 = h16(e13)) | (~ spl150_385 | ~ spl150_495)), inference(forward_demodulation, [], [f37740, f904])).
fof(f904, plain, (e22 = op2(e22, unit2)), inference(cnf_transformation, [], [f7])).
fof(f37740, plain, ((op2(e22, unit2) = h16(e13)) | (~ spl150_385 | ~ spl150_495)), inference(forward_demodulation, [], [f30970, f37665])).
fof(f37665, plain, ((unit2 = h18(e13)) | ~ spl150_495), inference(forward_demodulation, [], [f1588, f36013])).
fof(f36013, plain, ((op2(e22, e23) = unit2) | ~ spl150_495), inference(backward_demodulation, [], [f919, f4807])).
fof(f4807, plain, ((e22 = inv2(e23)) | ~ spl150_495), inference(avatar_component_clause, [], [f4805])).
fof(f30970, plain, ((h16(e13) = op2(e22, h18(e13))) | ~ spl150_385), inference(backward_demodulation, [], [f29176, f1576])).
fof(f1576, plain, (op2(e20, e23) = h16(e13)), inference(cnf_transformation, [], [f41])).
fof(f41, plain, ((op2(op2(e20, e23), op2(e20, e23)) = h16(e14)) & (op2(e20, e23) = h16(e13)) & (op2(e20, op2(op2(e20, e23), op2(e20, e23))) = h16(e12)) & (h16(e10) = op2(op2(e20, op2(op2(e20, e23), op2(e20, e23))), op2(e20, op2(op2(e20, e23), op2(e20, e23))))) & (e20 = h16(e15)) & (e23 = h16(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG027+1.p', ax41)).
fof(f29176, plain, ((op2(e20, e23) = op2(e22, h18(e13))) | ~ spl150_385), inference(forward_demodulation, [], [f5538, f4349])).
fof(f5538, plain, (op2(op2(e22, e22), e23) = op2(e22, h18(e13))), inference(backward_demodulation, [], [f770, f1588])).
fof(f770, plain, (op2(op2(e22, e22), e23) = op2(e22, op2(e22, e23))), inference(cnf_transformation, [], [f6])).
fof(f37309, plain, (~ spl150_656 | ~ spl150_770), inference(avatar_contradiction_clause, [], [f37308])).
fof(f37308, plain, ($false | (~ spl150_656 | ~ spl150_770)), inference(subsumption_resolution, [], [f37307, f5544])).
fof(f5544, plain, ~ (op2(e25, e23) = h18(e13)), inference(backward_demodulation, [], [f1282, f1588])).
fof(f1282, plain, ~ (op2(e22, e23) = op2(e25, e23)), inference(cnf_transformation, [], [f18])).
fof(f37307, plain, ((op2(e25, e23) = h18(e13)) | (~ spl150_656 | ~ spl150_770)), inference(forward_demodulation, [], [f20138, f7010])).
fof(f7010, plain, ((e23 = h19(e13)) | ~ spl150_770), inference(avatar_component_clause, [], [f7009])).
fof(f7009, plain, (spl150_770 <=> (e23 = h19(e13))), introduced(avatar_definition, [new_symbols(naming, [spl150_770])])).
fof(f20138, plain, ((h18(e13) = op2(e25, h19(e13))) | ~ spl150_656), inference(forward_demodulation, [], [f20118, f1588])).
fof(f20118, plain, ((op2(e22, e23) = op2(e25, h19(e13))) | ~ spl150_656), inference(backward_demodulation, [], [f5577, f6399])).
fof(f5577, plain, (op2(h10(e12), e23) = op2(e25, h19(e13))), inference(backward_demodulation, [], [f5346, f1594])).
fof(f5346, plain, (op2(e25, op2(e24, e23)) = op2(h10(e12), e23)), inference(backward_demodulation, [], [f890, f5339])).
fof(f890, plain, (op2(op2(e25, e24), e23) = op2(e25, op2(e24, e23))), inference(cnf_transformation, [], [f6])).
fof(f37230, plain, (~ spl150_9 | ~ spl150_28 | spl150_81 | ~ spl150_155), inference(avatar_contradiction_clause, [], [f37229])).
fof(f37229, plain, ($false | (~ spl150_9 | ~ spl150_28 | spl150_81 | ~ spl150_155)), inference(subsumption_resolution, [], [f37228, f3081])).
fof(f3081, plain, (~ (e12 = op1(e13, e14)) | spl150_81), inference(avatar_component_clause, [], [f3080])).
fof(f3080, plain, (spl150_81 <=> (e12 = op1(e13, e14))), introduced(avatar_definition, [new_symbols(naming, [spl150_81])])).
fof(f37228, plain, ((e12 = op1(e13, e14)) | (~ spl150_9 | ~ spl150_28 | ~ spl150_155)), inference(forward_demodulation, [], [f37227, f2782])).
fof(f37227, plain, ((op1(e13, e14) = op1(e15, e14)) | (~ spl150_28 | ~ spl150_155)), inference(forward_demodulation, [], [f19992, f3390])).
fof(f3390, plain, ((e14 = op1(e11, e14)) | ~ spl150_155), inference(avatar_component_clause, [], [f3388])).
fof(f3388, plain, (spl150_155 <=> (e14 = op1(e11, e14))), introduced(avatar_definition, [new_symbols(naming, [spl150_155])])).
fof(f19992, plain, ((op1(e13, e14) = op1(e15, op1(e11, e14))) | ~ spl150_28), inference(forward_demodulation, [], [f590, f2861])).
fof(f590, plain, (op1(op1(e15, e11), e14) = op1(e15, op1(e11, e14))), inference(cnf_transformation, [], [f2])).
fof(f37039, plain, (~ spl150_121 | ~ spl150_127), inference(avatar_split_clause, [], [f34705, f3272, f3247])).
fof(f3247, plain, (spl150_121 <=> (e10 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl150_121])])).
fof(f34705, plain, (~ (e10 = op1(e12, e13)) | ~ spl150_127), inference(forward_demodulation, [], [f1175, f3274])).
fof(f1175, plain, ~ (op1(e12, e12) = op1(e12, e13)), inference(cnf_transformation, [], [f17])).
fof(f17, plain, (~ (op1(e15, e14) = op1(e15, e15)) & ~ (op1(e15, e13) = op1(e15, e15)) & ~ (op1(e15, e13) = op1(e15, e14)) & ~ (op1(e15, e12) = op1(e15, e15)) & ~ (op1(e15, e12) = op1(e15, e14)) & ~ (op1(e15, e12) = op1(e15, e13)) & ~ (op1(e15, e11) = op1(e15, e15)) & ~ (op1(e15, e11) = op1(e15, e14)) & ~ (op1(e15, e11) = op1(e15, e13)) & ~ (op1(e15, e11) = op1(e15, e12)) & ~ (op1(e15, e10) = op1(e15, e15)) & ~ (op1(e15, e10) = op1(e15, e14)) & ~ (op1(e15, e10) = op1(e15, e13)) & ~ (op1(e15, e10) = op1(e15, e12)) & ~ (op1(e15, e10) = op1(e15, e11)) & ~ (op1(e14, e14) = op1(e14, e15)) & ~ (op1(e14, e13) = op1(e14, e15)) & ~ (op1(e14, e13) = op1(e14, e14)) & ~ (op1(e14, e12) = op1(e14, e15)) & ~ (op1(e14, e12) = op1(e14, e14)) & ~ (op1(e14, e12) = op1(e14, e13)) & ~ (op1(e14, e11) = op1(e14, e15)) & ~ (op1(e14, e11) = op1(e14, e14)) & ~ (op1(e14, e11) = op1(e14, e13)) & ~ (op1(e14, e11) = op1(e14, e12)) & ~ (op1(e14, e10) = op1(e14, e15)) & ~ (op1(e14, e10) = op1(e14, e14)) & ~ (op1(e14, e10) = op1(e14, e13)) & ~ (op1(e14, e10) = op1(e14, e12)) & ~ (op1(e14, e10) = op1(e14, e11)) & ~ (op1(e13, e14) = op1(e13, e15)) & ~ (op1(e13, e13) = op1(e13, e15)) & ~ (op1(e13, e13) = op1(e13, e14)) & ~ (op1(e13, e12) = op1(e13, e15)) & ~ (op1(e13, e12) = op1(e13, e14)) & ~ (op1(e13, e12) = op1(e13, e13)) & ~ (op1(e13, e11) = op1(e13, e15)) & ~ (op1(e13, e11) = op1(e13, e14)) & ~ (op1(e13, e11) = op1(e13, e13)) & ~ (op1(e13, e11) = op1(e13, e12)) & ~ (op1(e13, e10) = op1(e13, e15)) & ~ (op1(e13, e10) = op1(e13, e14)) & ~ (op1(e13, e10) = op1(e13, e13)) & ~ (op1(e13, e10) = op1(e13, e12)) & ~ (op1(e13, e10) = op1(e13, e11)) & ~ (op1(e12, e14) = op1(e12, e15)) & ~ (op1(e12, e13) = op1(e12, e15)) & ~ (op1(e12, e13) = op1(e12, e14)) & ~ (op1(e12, e12) = op1(e12, e15)) & ~ (op1(e12, e12) = op1(e12, e14)) & ~ (op1(e12, e12) = op1(e12, e13)) & ~ (op1(e12, e11) = op1(e12, e15)) & ~ (op1(e12, e11) = op1(e12, e14)) & ~ (op1(e12, e11) = op1(e12, e13)) & ~ (op1(e12, e11) = op1(e12, e12)) & ~ (op1(e12, e10) = op1(e12, e15)) & ~ (op1(e12, e10) = op1(e12, e14)) & ~ (op1(e12, e10) = op1(e12, e13)) & ~ (op1(e12, e10) = op1(e12, e12)) & ~ (op1(e12, e10) = op1(e12, e11)) & ~ (op1(e11, e14) = op1(e11, e15)) & ~ (op1(e11, e13) = op1(e11, e15)) & ~ (op1(e11, e13) = op1(e11, e14)) & ~ (op1(e11, e12) = op1(e11, e15)) & ~ (op1(e11, e12) = op1(e11, e14)) & ~ (op1(e11, e12) = op1(e11, e13)) & ~ (op1(e11, e11) = op1(e11, e15)) & ~ (op1(e11, e11) = op1(e11, e14)) & ~ (op1(e11, e11) = op1(e11, e13)) & ~ (op1(e11, e11) = op1(e11, e12)) & ~ (op1(e11, e10) = op1(e11, e15)) & ~ (op1(e11, e10) = op1(e11, e14)) & ~ (op1(e11, e10) = op1(e11, e13)) & ~ (op1(e11, e10) = op1(e11, e12)) & ~ (op1(e11, e10) = op1(e11, e11)) & ~ (op1(e10, e14) = op1(e10, e15)) & ~ (op1(e10, e13) = op1(e10, e15)) & ~ (op1(e10, e13) = op1(e10, e14)) & ~ (op1(e10, e12) = op1(e10, e15)) & ~ (op1(e10, e12) = op1(e10, e14)) & ~ (op1(e10, e12) = op1(e10, e13)) & ~ (op1(e10, e11) = op1(e10, e15)) & ~ (op1(e10, e11) = op1(e10, e14)) & ~ (op1(e10, e11) = op1(e10, e13)) & ~ (op1(e10, e11) = op1(e10, e12)) & ~ (op1(e10, e10) = op1(e10, e15)) & ~ (op1(e10, e10) = op1(e10, e14)) & ~ (op1(e10, e10) = op1(e10, e13)) & ~ (op1(e10, e10) = op1(e10, e12)) & ~ (op1(e10, e10) = op1(e10, e11)) & ~ (op1(e14, e15) = op1(e15, e15)) & ~ (op1(e13, e15) = op1(e15, e15)) & ~ (op1(e13, e15) = op1(e14, e15)) & ~ (op1(e12, e15) = op1(e15, e15)) & ~ (op1(e12, e15) = op1(e14, e15)) & ~ (op1(e12, e15) = op1(e13, e15)) & ~ (op1(e11, e15) = op1(e15, e15)) & ~ (op1(e11, e15) = op1(e14, e15)) & ~ (op1(e11, e15) = op1(e13, e15)) & ~ (op1(e11, e15) = op1(e12, e15)) & ~ (op1(e10, e15) = op1(e15, e15)) & ~ (op1(e10, e15) = op1(e14, e15)) & ~ (op1(e10, e15) = op1(e13, e15)) & ~ (op1(e10, e15) = op1(e12, e15)) & ~ (op1(e10, e15) = op1(e11, e15)) & ~ (op1(e14, e14) = op1(e15, e14)) & ~ (op1(e13, e14) = op1(e15, e14)) & ~ (op1(e13, e14) = op1(e14, e14)) & ~ (op1(e12, e14) = op1(e15, e14)) & ~ (op1(e12, e14) = op1(e14, e14)) & ~ (op1(e12, e14) = op1(e13, e14)) & ~ (op1(e11, e14) = op1(e15, e14)) & ~ (op1(e11, e14) = op1(e14, e14)) & ~ (op1(e11, e14) = op1(e13, e14)) & ~ (op1(e11, e14) = op1(e12, e14)) & ~ (op1(e10, e14) = op1(e15, e14)) & ~ (op1(e10, e14) = op1(e14, e14)) & ~ (op1(e10, e14) = op1(e13, e14)) & ~ (op1(e10, e14) = op1(e12, e14)) & ~ (op1(e10, e14) = op1(e11, e14)) & ~ (op1(e14, e13) = op1(e15, e13)) & ~ (op1(e13, e13) = op1(e15, e13)) & ~ (op1(e13, e13) = op1(e14, e13)) & ~ (op1(e12, e13) = op1(e15, e13)) & ~ (op1(e12, e13) = op1(e14, e13)) & ~ (op1(e12, e13) = op1(e13, e13)) & ~ (op1(e11, e13) = op1(e15, e13)) & ~ (op1(e11, e13) = op1(e14, e13)) & ~ (op1(e11, e13) = op1(e13, e13)) & ~ (op1(e11, e13) = op1(e12, e13)) & ~ (op1(e10, e13) = op1(e15, e13)) & ~ (op1(e10, e13) = op1(e14, e13)) & ~ (op1(e10, e13) = op1(e13, e13)) & ~ (op1(e10, e13) = op1(e12, e13)) & ~ (op1(e10, e13) = op1(e11, e13)) & ~ (op1(e14, e12) = op1(e15, e12)) & ~ (op1(e13, e12) = op1(e15, e12)) & ~ (op1(e13, e12) = op1(e14, e12)) & ~ (op1(e12, e12) = op1(e15, e12)) & ~ (op1(e12, e12) = op1(e14, e12)) & ~ (op1(e12, e12) = op1(e13, e12)) & ~ (op1(e11, e12) = op1(e15, e12)) & ~ (op1(e11, e12) = op1(e14, e12)) & ~ (op1(e11, e12) = op1(e13, e12)) & ~ (op1(e11, e12) = op1(e12, e12)) & ~ (op1(e10, e12) = op1(e15, e12)) & ~ (op1(e10, e12) = op1(e14, e12)) & ~ (op1(e10, e12) = op1(e13, e12)) & ~ (op1(e10, e12) = op1(e12, e12)) & ~ (op1(e10, e12) = op1(e11, e12)) & ~ (op1(e14, e11) = op1(e15, e11)) & ~ (op1(e13, e11) = op1(e15, e11)) & ~ (op1(e13, e11) = op1(e14, e11)) & ~ (op1(e12, e11) = op1(e15, e11)) & ~ (op1(e12, e11) = op1(e14, e11)) & ~ (op1(e12, e11) = op1(e13, e11)) & ~ (op1(e11, e11) = op1(e15, e11)) & ~ (op1(e11, e11) = op1(e14, e11)) & ~ (op1(e11, e11) = op1(e13, e11)) & ~ (op1(e11, e11) = op1(e12, e11)) & ~ (op1(e10, e11) = op1(e15, e11)) & ~ (op1(e10, e11) = op1(e14, e11)) & ~ (op1(e10, e11) = op1(e13, e11)) & ~ (op1(e10, e11) = op1(e12, e11)) & ~ (op1(e10, e11) = op1(e11, e11)) & ~ (op1(e14, e10) = op1(e15, e10)) & ~ (op1(e13, e10) = op1(e15, e10)) & ~ (op1(e13, e10) = op1(e14, e10)) & ~ (op1(e12, e10) = op1(e15, e10)) & ~ (op1(e12, e10) = op1(e14, e10)) & ~ (op1(e12, e10) = op1(e13, e10)) & ~ (op1(e11, e10) = op1(e15, e10)) & ~ (op1(e11, e10) = op1(e14, e10)) & ~ (op1(e11, e10) = op1(e13, e10)) & ~ (op1(e11, e10) = op1(e12, e10)) & ~ (op1(e10, e10) = op1(e15, e10)) & ~ (op1(e10, e10) = op1(e14, e10)) & ~ (op1(e10, e10) = op1(e13, e10)) & ~ (op1(e10, e10) = op1(e12, e10)) & ~ (op1(e10, e10) = op1(e11, e10))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG027+1.p', ax17)).
fof(f36634, plain, (~ spl150_122 | ~ spl150_127 | spl150_136 | ~ spl150_196), inference(avatar_contradiction_clause, [], [f36633])).
fof(f36633, plain, ($false | (~ spl150_122 | ~ spl150_127 | spl150_136 | ~ spl150_196)), inference(subsumption_resolution, [], [f36612, f3310])).
fof(f3310, plain, (~ (e13 = op1(e12, e11)) | spl150_136), inference(avatar_component_clause, [], [f3309])).
fof(f3309, plain, (spl150_136 <=> (e13 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl150_136])])).
fof(f36612, plain, ((e13 = op1(e12, e11)) | (~ spl150_122 | ~ spl150_127 | ~ spl150_196)), inference(backward_demodulation, [], [f34735, f3253])).
fof(f3253, plain, ((e11 = op1(e12, e13)) | ~ spl150_122), inference(avatar_component_clause, [], [f3251])).
fof(f3251, plain, (spl150_122 <=> (e11 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl150_122])])).
fof(f34735, plain, ((e13 = op1(e12, op1(e12, e13))) | (~ spl150_127 | ~ spl150_196)), inference(forward_demodulation, [], [f34734, f3561])).
fof(f34734, plain, ((op1(e10, e13) = op1(e12, op1(e12, e13))) | ~ spl150_127), inference(forward_demodulation, [], [f487, f3274])).
fof(f487, plain, (op1(op1(e12, e12), e13) = op1(e12, op1(e12, e13))), inference(cnf_transformation, [], [f2])).
fof(f36584, plain, (~ spl150_126 | ~ spl150_138), inference(avatar_split_clause, [], [f36566, f3317, f3267])).
fof(f3267, plain, (spl150_126 <=> (e15 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl150_126])])).
fof(f3317, plain, (spl150_138 <=> (e15 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl150_138])])).
fof(f36566, plain, (~ (e15 = op1(e12, e13)) | ~ spl150_138), inference(backward_demodulation, [], [f1172, f3319])).
fof(f3319, plain, ((e15 = op1(e12, e11)) | ~ spl150_138), inference(avatar_component_clause, [], [f3317])).
fof(f1172, plain, ~ (op1(e12, e11) = op1(e12, e13)), inference(cnf_transformation, [], [f17])).
fof(f36505, plain, (spl150_23 | ~ spl150_28 | ~ spl150_89 | ~ spl150_159), inference(avatar_split_clause, [], [f36493, f3405, f3113, f2859, f2838])).
fof(f36493, plain, ((e14 = op1(e15, e12)) | (~ spl150_28 | ~ spl150_89 | ~ spl150_159)), inference(backward_demodulation, [], [f30488, f3407])).
fof(f30488, plain, ((e14 = op1(e15, op1(e11, e13))) | (~ spl150_28 | ~ spl150_89)), inference(forward_demodulation, [], [f19998, f3115])).
fof(f19998, plain, ((op1(e13, e13) = op1(e15, op1(e11, e13))) | ~ spl150_28), inference(forward_demodulation, [], [f589, f2861])).
fof(f589, plain, (op1(op1(e15, e11), e13) = op1(e15, op1(e11, e13))), inference(cnf_transformation, [], [f2])).
fof(f36097, plain, (~ spl150_385 | ~ spl150_477), inference(avatar_contradiction_clause, [], [f36096])).
fof(f36096, plain, ($false | (~ spl150_385 | ~ spl150_477)), inference(subsumption_resolution, [], [f36095, f1422])).
fof(f1422, plain, ~ (e20 = e22), inference(cnf_transformation, [], [f20])).
fof(f20, plain, (~ (e24 = e25) & ~ (e23 = e25) & ~ (e23 = e24) & ~ (e22 = e25) & ~ (e22 = e24) & ~ (e22 = e23) & ~ (e21 = e25) & ~ (e21 = e24) & ~ (e21 = e23) & ~ (e21 = e22) & ~ (e20 = e25) & ~ (e20 = e24) & ~ (e20 = e23) & ~ (e20 = e22) & ~ (e20 = e21)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG027+1.p', ax20)).
fof(f36095, plain, ((e20 = e22) | (~ spl150_385 | ~ spl150_477)), inference(forward_demodulation, [], [f36041, f4349])).
fof(f36041, plain, ((e22 = op2(e22, e22)) | ~ spl150_477), inference(backward_demodulation, [], [f903, f4732])).
fof(f4732, plain, ((e22 = unit2) | ~ spl150_477), inference(avatar_component_clause, [], [f4730])).
fof(f4730, plain, (spl150_477 <=> (e22 = unit2)), introduced(avatar_definition, [new_symbols(naming, [spl150_477])])).
fof(f903, plain, (e22 = op2(unit2, e22)), inference(cnf_transformation, [], [f7])).
fof(f35108, plain, (~ spl150_385 | ~ spl150_489 | ~ spl150_656), inference(avatar_contradiction_clause, [], [f35107])).
fof(f35107, plain, ($false | (~ spl150_385 | ~ spl150_489 | ~ spl150_656)), inference(subsumption_resolution, [], [f35106, f1425])).
fof(f1425, plain, ~ (e20 = e25), inference(cnf_transformation, [], [f20])).
fof(f35106, plain, ((e20 = e25) | (~ spl150_385 | ~ spl150_489 | ~ spl150_656)), inference(forward_demodulation, [], [f910, f34844])).
fof(f34844, plain, ((e20 = op2(e25, unit2)) | (~ spl150_385 | ~ spl150_489 | ~ spl150_656)), inference(backward_demodulation, [], [f26890, f34834])).
fof(f34834, plain, ((unit2 = h14(e13)) | ~ spl150_489), inference(forward_demodulation, [], [f1564, f33634])).
fof(f33634, plain, ((op2(e24, e22) = unit2) | ~ spl150_489), inference(backward_demodulation, [], [f920, f4782])).
fof(f4782, plain, ((e22 = inv2(e24)) | ~ spl150_489), inference(avatar_component_clause, [], [f4780])).
fof(f4780, plain, (spl150_489 <=> (e22 = inv2(e24))), introduced(avatar_definition, [new_symbols(naming, [spl150_489])])).
fof(f920, plain, (unit2 = op2(e24, inv2(e24))), inference(cnf_transformation, [], [f8])).
fof(f1564, plain, (op2(e24, e22) = h14(e13)), inference(cnf_transformation, [], [f39])).
fof(f39, plain, ((op2(op2(e24, e22), op2(e24, e22)) = h14(e14)) & (op2(e24, e22) = h14(e13)) & (op2(e24, op2(op2(e24, e22), op2(e24, e22))) = h14(e12)) & (h14(e10) = op2(op2(e24, op2(op2(e24, e22), op2(e24, e22))), op2(e24, op2(op2(e24, e22), op2(e24, e22))))) & (e24 = h14(e15)) & (e22 = h14(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG027+1.p', ax39)).
fof(f26890, plain, ((e20 = op2(e25, h14(e13))) | (~ spl150_385 | ~ spl150_656)), inference(forward_demodulation, [], [f20117, f4349])).
fof(f20117, plain, ((op2(e22, e22) = op2(e25, h14(e13))) | ~ spl150_656), inference(backward_demodulation, [], [f5460, f6399])).
fof(f5460, plain, (op2(h10(e12), e22) = op2(e25, h14(e13))), inference(backward_demodulation, [], [f5345, f1564])).
fof(f5345, plain, (op2(e25, op2(e24, e22)) = op2(h10(e12), e22)), inference(backward_demodulation, [], [f889, f5339])).
fof(f889, plain, (op2(op2(e25, e24), e22) = op2(e25, op2(e24, e22))), inference(cnf_transformation, [], [f6])).
fof(f910, plain, (e25 = op2(e25, unit2)), inference(cnf_transformation, [], [f7])).
fof(f35096, plain, (~ spl150_347 | spl150_375 | ~ spl150_385 | ~ spl150_459 | ~ spl150_840), inference(avatar_contradiction_clause, [], [f35095])).
fof(f35095, plain, ($false | (~ spl150_347 | spl150_375 | ~ spl150_385 | ~ spl150_459 | ~ spl150_840)), inference(subsumption_resolution, [], [f35094, f4306])).
fof(f4306, plain, (~ (e22 = op2(e22, e24)) | spl150_375), inference(avatar_component_clause, [], [f4305])).
fof(f4305, plain, (spl150_375 <=> (e22 = op2(e22, e24))), introduced(avatar_definition, [new_symbols(naming, [spl150_375])])).
fof(f35094, plain, ((e22 = op2(e22, e24)) | (~ spl150_347 | ~ spl150_385 | ~ spl150_459 | ~ spl150_840)), inference(forward_demodulation, [], [f35093, f33712])).
fof(f33712, plain, ((e22 = h2(e13)) | (~ spl150_385 | ~ spl150_459)), inference(backward_demodulation, [], [f27140, f4657])).
fof(f27140, plain, ((op2(e20, e22) = h2(e13)) | ~ spl150_385), inference(forward_demodulation, [], [f27139, f1492])).
fof(f1492, plain, (op2(e22, e20) = h2(e13)), inference(cnf_transformation, [], [f27])).
fof(f27, plain, ((op2(op2(e22, e20), op2(e22, e20)) = h2(e14)) & (op2(e22, e20) = h2(e13)) & (op2(e22, op2(op2(e22, e20), op2(e22, e20))) = h2(e12)) & (h2(e10) = op2(op2(e22, op2(op2(e22, e20), op2(e22, e20))), op2(e22, op2(op2(e22, e20), op2(e22, e20))))) & (e22 = h2(e15)) & (e20 = h2(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG027+1.p', ax27)).
fof(f27139, plain, ((op2(e20, e22) = op2(e22, e20)) | ~ spl150_385), inference(forward_demodulation, [], [f769, f4349])).
fof(f769, plain, (op2(op2(e22, e22), e22) = op2(e22, op2(e22, e22))), inference(cnf_transformation, [], [f6])).
fof(f35093, plain, ((e22 = op2(h2(e13), e24)) | (~ spl150_347 | ~ spl150_385 | ~ spl150_840)), inference(forward_demodulation, [], [f35092, f35060])).
fof(f35060, plain, ((e22 = op2(e22, h18(e13))) | (~ spl150_385 | ~ spl150_840)), inference(forward_demodulation, [], [f30970, f7380])).
fof(f7380, plain, ((e22 = h16(e13)) | ~ spl150_840), inference(avatar_component_clause, [], [f7379])).
fof(f35092, plain, ((op2(h2(e13), e24) = op2(e22, h18(e13))) | (~ spl150_347 | ~ spl150_840)), inference(forward_demodulation, [], [f5615, f35032])).
fof(f35032, plain, ((h18(e13) = h21(e13)) | (~ spl150_347 | ~ spl150_840)), inference(forward_demodulation, [], [f35031, f1588])).
fof(f35031, plain, ((op2(e22, e23) = h21(e13)) | (~ spl150_347 | ~ spl150_840)), inference(forward_demodulation, [], [f26924, f7380])).
fof(f26924, plain, ((h21(e13) = op2(h16(e13), e23)) | ~ spl150_347), inference(backward_demodulation, [], [f26733, f1606])).
fof(f26733, plain, ((op2(e20, e24) = op2(h16(e13), e23)) | ~ spl150_347), inference(forward_demodulation, [], [f5488, f4190])).
fof(f5488, plain, (op2(e20, op2(e23, e23)) = op2(h16(e13), e23)), inference(backward_demodulation, [], [f704, f1576])).
fof(f704, plain, (op2(op2(e20, e23), e23) = op2(e20, op2(e23, e23))), inference(cnf_transformation, [], [f6])).
fof(f5615, plain, (op2(h2(e13), e24) = op2(e22, h21(e13))), inference(backward_demodulation, [], [f5120, f1606])).
fof(f5120, plain, (op2(e22, op2(e20, e24)) = op2(h2(e13), e24)), inference(backward_demodulation, [], [f759, f1492])).
fof(f759, plain, (op2(op2(e22, e20), e24) = op2(e22, op2(e20, e24))), inference(cnf_transformation, [], [f6])).
fof(f34787, plain, (spl150_253 | ~ spl150_217), inference(avatar_split_clause, [], [f33914, f3647, f3797])).
fof(f3797, plain, (spl150_253 <=> (e10 = inv1(e10))), introduced(avatar_definition, [new_symbols(naming, [spl150_253])])).
fof(f33914, plain, ((e10 = inv1(e10)) | ~ spl150_217), inference(backward_demodulation, [], [f930, f3649])).
fof(f930, plain, (unit1 = inv1(unit1)), inference(cnf_transformation, [], [f9])).
fof(f9, plain, (unit1 = inv1(unit1)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG027+1.p', ax9)).
fof(f34619, plain, (spl150_138 | ~ spl150_9 | ~ spl150_36 | ~ spl150_217 | ~ spl150_251), inference(avatar_split_clause, [], [f34618, f3788, f3647, f2892, f2780, f3317])).
fof(f3788, plain, (spl150_251 <=> (e14 = inv1(e11))), introduced(avatar_definition, [new_symbols(naming, [spl150_251])])).
fof(f34618, plain, ((e15 = op1(e12, e11)) | (~ spl150_9 | ~ spl150_36 | ~ spl150_217 | ~ spl150_251)), inference(forward_demodulation, [], [f34617, f2894])).
fof(f34617, plain, ((op1(e12, e11) = op1(e15, e10)) | (~ spl150_9 | ~ spl150_217 | ~ spl150_251)), inference(forward_demodulation, [], [f30381, f3649])).
fof(f30381, plain, ((op1(e12, e11) = op1(e15, unit1)) | (~ spl150_9 | ~ spl150_251)), inference(forward_demodulation, [], [f19932, f27722])).
fof(f27722, plain, ((op1(e14, e11) = unit1) | ~ spl150_251), inference(backward_demodulation, [], [f632, f3790])).
fof(f3790, plain, ((e14 = inv1(e11)) | ~ spl150_251), inference(avatar_component_clause, [], [f3788])).
fof(f19932, plain, ((op1(e12, e11) = op1(e15, op1(e14, e11))) | ~ spl150_9), inference(forward_demodulation, [], [f605, f2782])).
fof(f605, plain, (op1(op1(e15, e14), e11) = op1(e15, op1(e14, e11))), inference(cnf_transformation, [], [f2])).
fof(f34503, plain, (~ spl150_25 | ~ spl150_28), inference(avatar_contradiction_clause, [], [f34502])).
fof(f34502, plain, ($false | (~ spl150_25 | ~ spl150_28)), inference(subsumption_resolution, [], [f34501, f1408])).
fof(f1408, plain, ~ (e10 = e13), inference(cnf_transformation, [], [f19])).
fof(f19, plain, (~ (e14 = e15) & ~ (e13 = e15) & ~ (e13 = e14) & ~ (e12 = e15) & ~ (e12 = e14) & ~ (e12 = e13) & ~ (e11 = e15) & ~ (e11 = e14) & ~ (e11 = e13) & ~ (e11 = e12) & ~ (e10 = e15) & ~ (e10 = e14) & ~ (e10 = e13) & ~ (e10 = e12) & ~ (e10 = e11)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG027+1.p', ax19)).
fof(f34501, plain, ((e10 = e13) | (~ spl150_25 | ~ spl150_28)), inference(backward_demodulation, [], [f2861, f2849])).
fof(f2849, plain, ((e10 = op1(e15, e11)) | ~ spl150_25), inference(avatar_component_clause, [], [f2847])).
fof(f34488, plain, (~ spl150_9 | spl150_34 | ~ spl150_37 | ~ spl150_126 | ~ spl150_127 | ~ spl150_196), inference(avatar_contradiction_clause, [], [f34487])).
fof(f34487, plain, ($false | (~ spl150_9 | spl150_34 | ~ spl150_37 | ~ spl150_126 | ~ spl150_127 | ~ spl150_196)), inference(subsumption_resolution, [], [f34466, f2885])).
fof(f2885, plain, (~ (e13 = op1(e15, e10)) | spl150_34), inference(avatar_component_clause, [], [f2884])).
fof(f2884, plain, (spl150_34 <=> (e13 = op1(e15, e10))), introduced(avatar_definition, [new_symbols(naming, [spl150_34])])).
fof(f34466, plain, ((e13 = op1(e15, e10)) | (~ spl150_9 | ~ spl150_37 | ~ spl150_126 | ~ spl150_127 | ~ spl150_196)), inference(backward_demodulation, [], [f34075, f2899])).
fof(f2899, plain, ((e10 = op1(e14, e15)) | ~ spl150_37), inference(avatar_component_clause, [], [f2897])).
fof(f2897, plain, (spl150_37 <=> (e10 = op1(e14, e15))), introduced(avatar_definition, [new_symbols(naming, [spl150_37])])).
fof(f34075, plain, ((e13 = op1(e15, op1(e14, e15))) | (~ spl150_9 | ~ spl150_126 | ~ spl150_127 | ~ spl150_196)), inference(backward_demodulation, [], [f32732, f3561])).
fof(f32732, plain, ((op1(e10, e13) = op1(e15, op1(e14, e15))) | (~ spl150_9 | ~ spl150_126 | ~ spl150_127)), inference(backward_demodulation, [], [f19919, f32722])).
fof(f32722, plain, ((op1(e10, e13) = op1(e12, e15)) | (~ spl150_126 | ~ spl150_127)), inference(forward_demodulation, [], [f32721, f3274])).
fof(f32721, plain, ((op1(e12, e15) = op1(op1(e12, e12), e13)) | ~ spl150_126), inference(forward_demodulation, [], [f487, f3269])).
fof(f3269, plain, ((e15 = op1(e12, e13)) | ~ spl150_126), inference(avatar_component_clause, [], [f3267])).
fof(f19919, plain, ((op1(e12, e15) = op1(e15, op1(e14, e15))) | ~ spl150_9), inference(forward_demodulation, [], [f609, f2782])).
fof(f609, plain, (op1(op1(e15, e14), e15) = op1(e15, op1(e14, e15))), inference(cnf_transformation, [], [f2])).
fof(f34121, plain, (spl150_112 | ~ spl150_126 | ~ spl150_127 | ~ spl150_141 | ~ spl150_186 | ~ spl150_196), inference(avatar_split_clause, [], [f34113, f3559, f3517, f3330, f3272, f3267, f3209])).
fof(f34113, plain, ((e13 = op1(e12, e15)) | (~ spl150_126 | ~ spl150_127 | ~ spl150_141 | ~ spl150_186 | ~ spl150_196)), inference(backward_demodulation, [], [f34079, f3519])).
fof(f34079, plain, ((e13 = op1(e12, op1(e10, e15))) | (~ spl150_126 | ~ spl150_127 | ~ spl150_141 | ~ spl150_196)), inference(backward_demodulation, [], [f32741, f3561])).
fof(f32741, plain, ((op1(e10, e13) = op1(e12, op1(e10, e15))) | (~ spl150_126 | ~ spl150_127 | ~ spl150_141)), inference(backward_demodulation, [], [f32586, f32722])).
fof(f32586, plain, ((op1(e12, e15) = op1(e12, op1(e10, e15))) | ~ spl150_141), inference(forward_demodulation, [], [f477, f3332])).
fof(f477, plain, (op1(op1(e12, e10), e15) = op1(e12, op1(e10, e15))), inference(cnf_transformation, [], [f2])).
fof(f34096, plain, (~ spl150_28 | spl150_102 | ~ spl150_126 | ~ spl150_127 | ~ spl150_196), inference(avatar_contradiction_clause, [], [f34095])).
fof(f34095, plain, ($false | (~ spl150_28 | spl150_102 | ~ spl150_126 | ~ spl150_127 | ~ spl150_196)), inference(subsumption_resolution, [], [f34078, f3168])).
fof(f3168, plain, (~ (e15 = op1(e13, e11)) | spl150_102), inference(avatar_component_clause, [], [f3167])).
fof(f34078, plain, ((e15 = op1(e13, e11)) | (~ spl150_28 | ~ spl150_126 | ~ spl150_127 | ~ spl150_196)), inference(backward_demodulation, [], [f32738, f3561])).
fof(f32738, plain, ((e15 = op1(op1(e10, e13), e11)) | (~ spl150_28 | ~ spl150_126 | ~ spl150_127)), inference(backward_demodulation, [], [f30440, f32722])).
fof(f30440, plain, ((e15 = op1(op1(e12, e15), e11)) | (~ spl150_28 | ~ spl150_126)), inference(forward_demodulation, [], [f19963, f3269])).
fof(f19963, plain, ((op1(e12, e13) = op1(op1(e12, e15), e11)) | ~ spl150_28), inference(forward_demodulation, [], [f503, f2861])).
fof(f503, plain, (op1(op1(e12, e15), e11) = op1(e12, op1(e15, e11))), inference(cnf_transformation, [], [f2])).
fof(f34091, plain, (~ spl150_9 | spl150_79 | ~ spl150_126 | ~ spl150_127 | ~ spl150_196), inference(avatar_contradiction_clause, [], [f34090])).
fof(f34090, plain, ($false | (~ spl150_9 | spl150_79 | ~ spl150_126 | ~ spl150_127 | ~ spl150_196)), inference(subsumption_resolution, [], [f34076, f3073])).
fof(f3073, plain, (~ (e10 = op1(e13, e14)) | spl150_79), inference(avatar_component_clause, [], [f3072])).
fof(f34076, plain, ((e10 = op1(e13, e14)) | (~ spl150_9 | ~ spl150_126 | ~ spl150_127 | ~ spl150_196)), inference(backward_demodulation, [], [f32735, f3561])).
fof(f32735, plain, ((e10 = op1(op1(e10, e13), e14)) | (~ spl150_9 | ~ spl150_126 | ~ spl150_127)), inference(backward_demodulation, [], [f25810, f32722])).
fof(f25810, plain, ((e10 = op1(op1(e12, e15), e14)) | (~ spl150_9 | ~ spl150_127)), inference(forward_demodulation, [], [f19961, f3274])).
fof(f19961, plain, ((op1(e12, e12) = op1(op1(e12, e15), e14)) | ~ spl150_9), inference(forward_demodulation, [], [f506, f2782])).
fof(f506, plain, (op1(op1(e12, e15), e14) = op1(e12, op1(e15, e14))), inference(cnf_transformation, [], [f2])).
fof(f34086, plain, (~ spl150_160 | ~ spl150_196), inference(avatar_split_clause, [], [f34060, f3559, f3409])).
fof(f3409, plain, (spl150_160 <=> (e13 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl150_160])])).
fof(f34060, plain, (~ (e13 = op1(e11, e13)) | ~ spl150_196), inference(backward_demodulation, [], [f1091, f3561])).
fof(f1091, plain, ~ (op1(e10, e13) = op1(e11, e13)), inference(cnf_transformation, [], [f17])).
fof(f33961, plain, (spl150_211 | ~ spl150_217), inference(avatar_split_clause, [], [f33912, f3647, f3622])).
fof(f33912, plain, ((e10 = op1(e10, e10)) | ~ spl150_217), inference(backward_demodulation, [], [f617, f3649])).
fof(f617, plain, (e10 = op1(e10, unit1)), inference(cnf_transformation, [], [f3])).
fof(f3, plain, (((e15 = unit1) | (e14 = unit1) | (e13 = unit1) | (e12 = unit1) | (e11 = unit1) | (e10 = unit1)) & (e15 = op1(e15, unit1)) & (e15 = op1(unit1, e15)) & (e14 = op1(e14, unit1)) & (e14 = op1(unit1, e14)) & (e13 = op1(e13, unit1)) & (e13 = op1(unit1, e13)) & (e12 = op1(e12, unit1)) & (e12 = op1(unit1, e12)) & (e11 = op1(e11, unit1)) & (e11 = op1(unit1, e11)) & (e10 = op1(e10, unit1)) & (e10 = op1(unit1, e10))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG027+1.p', ax3)).
fof(f33632, plain, (~ spl150_347 | spl150_479 | ~ spl150_496), inference(avatar_contradiction_clause, [], [f33631])).
fof(f33631, plain, ($false | (~ spl150_347 | spl150_479 | ~ spl150_496)), inference(subsumption_resolution, [], [f33630, f4739])).
fof(f4739, plain, (~ (e24 = unit2) | spl150_479), inference(avatar_component_clause, [], [f4738])).
fof(f4738, plain, (spl150_479 <=> (e24 = unit2)), introduced(avatar_definition, [new_symbols(naming, [spl150_479])])).
fof(f33630, plain, ((e24 = unit2) | (~ spl150_347 | ~ spl150_496)), inference(forward_demodulation, [], [f33622, f4190])).
fof(f33622, plain, ((op2(e23, e23) = unit2) | ~ spl150_496), inference(backward_demodulation, [], [f919, f4811])).
fof(f4811, plain, ((e23 = inv2(e23)) | ~ spl150_496), inference(avatar_component_clause, [], [f4809])).
fof(f4809, plain, (spl150_496 <=> (e23 = inv2(e23))), introduced(avatar_definition, [new_symbols(naming, [spl150_496])])).
fof(f33619, plain, (spl150_489 | ~ spl150_503), inference(avatar_split_clause, [], [f33610, f4838, f4780])).
fof(f4838, plain, (spl150_503 <=> (e24 = inv2(e22))), introduced(avatar_definition, [new_symbols(naming, [spl150_503])])).
fof(f33610, plain, ((e22 = inv2(e24)) | ~ spl150_503), inference(backward_demodulation, [], [f991, f4840])).
fof(f4840, plain, ((e24 = inv2(e22)) | ~ spl150_503), inference(avatar_component_clause, [], [f4838])).
fof(f991, plain, (e22 = inv2(inv2(e22))), inference(cnf_transformation, [], [f14])).
fof(f14, plain, ((e25 = inv2(inv2(e25))) & (e24 = inv2(inv2(e24))) & (e23 = inv2(inv2(e23))) & (e22 = inv2(inv2(e22))) & (e21 = inv2(inv2(e21))) & (e20 = inv2(inv2(e20)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG027+1.p', ax14)).
fof(f33091, plain, (spl150_770 | ~ spl150_340 | ~ spl150_347), inference(avatar_split_clause, [], [f33090, f4188, f4159, f7009])).
fof(f4159, plain, (spl150_340 <=> (e23 = op2(e23, e24))), introduced(avatar_definition, [new_symbols(naming, [spl150_340])])).
fof(f33090, plain, ((e23 = h19(e13)) | (~ spl150_340 | ~ spl150_347)), inference(forward_demodulation, [], [f26724, f4161])).
fof(f4161, plain, ((e23 = op2(e23, e24)) | ~ spl150_340), inference(avatar_component_clause, [], [f4159])).
fof(f33060, plain, (spl150_478 | ~ spl150_286 | ~ spl150_510), inference(avatar_split_clause, [], [f31621, f4867, f3934, f4734])).
fof(f4734, plain, (spl150_478 <=> (e23 = unit2)), introduced(avatar_definition, [new_symbols(naming, [spl150_478])])).
fof(f4867, plain, (spl150_510 <=> (e25 = inv2(e21))), introduced(avatar_definition, [new_symbols(naming, [spl150_510])])).
fof(f31621, plain, ((e23 = unit2) | (~ spl150_286 | ~ spl150_510)), inference(forward_demodulation, [], [f31613, f3936])).
fof(f31613, plain, ((op2(e25, e21) = unit2) | ~ spl150_510), inference(backward_demodulation, [], [f915, f4869])).
fof(f4869, plain, ((e25 = inv2(e21)) | ~ spl150_510), inference(avatar_component_clause, [], [f4867])).
fof(f33059, plain, (spl150_478 | ~ spl150_286 | ~ spl150_482), inference(avatar_split_clause, [], [f31647, f4751, f3934, f4734])).
fof(f4751, plain, (spl150_482 <=> (e21 = inv2(e25))), introduced(avatar_definition, [new_symbols(naming, [spl150_482])])).
fof(f31647, plain, ((e23 = unit2) | (~ spl150_286 | ~ spl150_482)), inference(forward_demodulation, [], [f31640, f3936])).
fof(f31640, plain, ((op2(e25, e21) = unit2) | ~ spl150_482), inference(backward_demodulation, [], [f922, f4753])).
fof(f4753, plain, ((e21 = inv2(e25)) | ~ spl150_482), inference(avatar_component_clause, [], [f4751])).
fof(f922, plain, (unit2 = op2(e25, inv2(e25))), inference(cnf_transformation, [], [f8])).
fof(f32825, plain, (~ spl150_1655 | ~ spl150_126 | ~ spl150_656 | ~ spl150_954 | spl150_1938), inference(avatar_split_clause, [], [f32824, f15129, f7985, f6398, f3267, f12977])).
fof(f15129, plain, (spl150_1938 <=> (h10(op1(e12, e13)) = op2(h10(e12), h10(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl150_1938])])).
fof(f32824, plain, (~ (e25 = h18(e13)) | (~ spl150_126 | ~ spl150_656 | ~ spl150_954 | spl150_1938)), inference(forward_demodulation, [], [f32823, f1537])).
fof(f32823, plain, (~ (h10(e15) = h18(e13)) | (~ spl150_126 | ~ spl150_656 | ~ spl150_954 | spl150_1938)), inference(forward_demodulation, [], [f20147, f3269])).
fof(f20147, plain, (~ (h18(e13) = h10(op1(e12, e13))) | (~ spl150_656 | ~ spl150_954 | spl150_1938)), inference(forward_demodulation, [], [f20128, f1588])).
fof(f20128, plain, (~ (op2(e22, e23) = h10(op1(e12, e13))) | (~ spl150_656 | ~ spl150_954 | spl150_1938)), inference(backward_demodulation, [], [f20087, f6399])).
fof(f20087, plain, (~ (h10(op1(e12, e13)) = op2(h10(e12), e23)) | (~ spl150_954 | spl150_1938)), inference(backward_demodulation, [], [f15131, f7986])).
fof(f15131, plain, (~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | spl150_1938), inference(avatar_component_clause, [], [f15129])).
fof(f32822, plain, (~ spl150_162 | ~ spl150_126), inference(avatar_split_clause, [], [f32821, f3267, f3417])).
fof(f3417, plain, (spl150_162 <=> (e15 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl150_162])])).
fof(f32821, plain, (~ (e15 = op1(e11, e13)) | ~ spl150_126), inference(forward_demodulation, [], [f1096, f3269])).
fof(f1096, plain, ~ (op1(e11, e13) = op1(e12, e13)), inference(cnf_transformation, [], [f17])).
fof(f32585, plain, (~ spl150_154 | ~ spl150_46), inference(avatar_split_clause, [], [f30045, f2934, f3384])).
fof(f3384, plain, (spl150_154 <=> (e13 = op1(e11, e14))), introduced(avatar_definition, [new_symbols(naming, [spl150_154])])).
fof(f30045, plain, (~ (e13 = op1(e11, e14)) | ~ spl150_46), inference(forward_demodulation, [], [f1113, f2936])).
fof(f1113, plain, ~ (op1(e11, e14) = op1(e14, e14)), inference(cnf_transformation, [], [f17])).
fof(f32545, plain, (~ spl150_135 | ~ spl150_141), inference(avatar_split_clause, [], [f30455, f3330, f3305])).
fof(f3305, plain, (spl150_135 <=> (e12 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl150_135])])).
fof(f30455, plain, (~ (e12 = op1(e12, e11)) | ~ spl150_141), inference(forward_demodulation, [], [f1166, f3332])).
fof(f1166, plain, ~ (op1(e12, e10) = op1(e12, e11)), inference(cnf_transformation, [], [f17])).
fof(f32411, plain, (~ spl150_23 | ~ spl150_39 | ~ spl150_46 | ~ spl150_127), inference(avatar_contradiction_clause, [], [f32410])).
fof(f32410, plain, ($false | (~ spl150_23 | ~ spl150_39 | ~ spl150_46 | ~ spl150_127)), inference(subsumption_resolution, [], [f32409, f1408])).
fof(f32409, plain, ((e10 = e13) | (~ spl150_23 | ~ spl150_39 | ~ spl150_46 | ~ spl150_127)), inference(forward_demodulation, [], [f32393, f3274])).
fof(f32393, plain, ((e13 = op1(e12, e12)) | (~ spl150_23 | ~ spl150_39 | ~ spl150_46)), inference(backward_demodulation, [], [f30126, f2907])).
fof(f2907, plain, ((e12 = op1(e14, e15)) | ~ spl150_39), inference(avatar_component_clause, [], [f2905])).
fof(f2905, plain, (spl150_39 <=> (e12 = op1(e14, e15))), introduced(avatar_definition, [new_symbols(naming, [spl150_39])])).
fof(f30126, plain, ((e13 = op1(op1(e14, e15), e12)) | (~ spl150_23 | ~ spl150_46)), inference(forward_demodulation, [], [f23868, f2936])).
fof(f23868, plain, ((op1(e14, e14) = op1(op1(e14, e15), e12)) | ~ spl150_23), inference(forward_demodulation, [], [f576, f2840])).
fof(f576, plain, (op1(op1(e14, e15), e12) = op1(e14, op1(e15, e12))), inference(cnf_transformation, [], [f2])).
fof(f32366, plain, (spl150_1 | ~ spl150_75 | ~ spl150_126 | ~ spl150_127), inference(avatar_split_clause, [], [f32365, f3272, f3267, f3055, f2747])).
fof(f32365, plain, ((e10 = op1(e15, e15)) | (~ spl150_75 | ~ spl150_126 | ~ spl150_127)), inference(forward_demodulation, [], [f32345, f3274])).
fof(f32345, plain, ((op1(e12, e12) = op1(e15, e15)) | (~ spl150_75 | ~ spl150_126)), inference(backward_demodulation, [], [f30312, f3057])).
fof(f30312, plain, ((op1(e15, e15) = op1(e12, op1(e13, e15))) | ~ spl150_126), inference(forward_demodulation, [], [f495, f3269])).
fof(f495, plain, (op1(op1(e12, e13), e15) = op1(e12, op1(e13, e15))), inference(cnf_transformation, [], [f2])).
fof(f32097, plain, (spl150_83 | ~ spl150_220), inference(avatar_contradiction_clause, [], [f32096])).
fof(f32096, plain, ($false | (spl150_83 | ~ spl150_220)), inference(subsumption_resolution, [], [f32058, f3089])).
fof(f3089, plain, (~ (e14 = op1(e13, e14)) | spl150_83), inference(avatar_component_clause, [], [f3088])).
fof(f3088, plain, (spl150_83 <=> (e14 = op1(e13, e14))), introduced(avatar_definition, [new_symbols(naming, [spl150_83])])).
fof(f32058, plain, ((e14 = op1(e13, e14)) | ~ spl150_220), inference(backward_demodulation, [], [f624, f3661])).
fof(f3661, plain, ((e13 = unit1) | ~ spl150_220), inference(avatar_component_clause, [], [f3659])).
fof(f3659, plain, (spl150_220 <=> (e13 = unit1)), introduced(avatar_definition, [new_symbols(naming, [spl150_220])])).
fof(f624, plain, (e14 = op1(unit1, e14)), inference(cnf_transformation, [], [f3])).
fof(f31704, plain, (spl150_279 | ~ spl150_480), inference(avatar_contradiction_clause, [], [f31703])).
fof(f31703, plain, ($false | (spl150_279 | ~ spl150_480)), inference(subsumption_resolution, [], [f31654, f3906])).
fof(f3906, plain, (~ (e22 = op2(e25, e22)) | spl150_279), inference(avatar_component_clause, [], [f3905])).
fof(f3905, plain, (spl150_279 <=> (e22 = op2(e25, e22))), introduced(avatar_definition, [new_symbols(naming, [spl150_279])])).
fof(f31654, plain, ((e22 = op2(e25, e22)) | ~ spl150_480), inference(backward_demodulation, [], [f903, f4744])).
fof(f4744, plain, ((e25 = unit2) | ~ spl150_480), inference(avatar_component_clause, [], [f4742])).
fof(f4742, plain, (spl150_480 <=> (e25 = unit2)), introduced(avatar_definition, [new_symbols(naming, [spl150_480])])).
fof(f31026, plain, (spl150_1655 | ~ spl150_497 | ~ spl150_656), inference(avatar_split_clause, [], [f31025, f6398, f4813, f12977])).
fof(f31025, plain, ((e25 = h18(e13)) | (~ spl150_497 | ~ spl150_656)), inference(backward_demodulation, [], [f30919, f910])).
fof(f30919, plain, ((op2(e25, unit2) = h18(e13)) | (~ spl150_497 | ~ spl150_656)), inference(forward_demodulation, [], [f20138, f30725])).
fof(f30725, plain, ((unit2 = h19(e13)) | ~ spl150_497), inference(backward_demodulation, [], [f1594, f27294])).
fof(f27294, plain, ((op2(e24, e23) = unit2) | ~ spl150_497), inference(backward_demodulation, [], [f919, f4815])).
fof(f30566, plain, (~ spl150_9 | spl150_15 | ~ spl150_18 | ~ spl150_89), inference(avatar_contradiction_clause, [], [f30565])).
fof(f30565, plain, ($false | (~ spl150_9 | spl150_15 | ~ spl150_18 | ~ spl150_89)), inference(subsumption_resolution, [], [f30564, f2806])).
fof(f2806, plain, (~ (e12 = op1(e15, e13)) | spl150_15), inference(avatar_component_clause, [], [f2805])).
fof(f2805, plain, (spl150_15 <=> (e12 = op1(e15, e13))), introduced(avatar_definition, [new_symbols(naming, [spl150_15])])).
fof(f30564, plain, ((e12 = op1(e15, e13)) | (~ spl150_9 | ~ spl150_18 | ~ spl150_89)), inference(forward_demodulation, [], [f25819, f2819])).
fof(f2819, plain, ((e15 = op1(e15, e13)) | ~ spl150_18), inference(avatar_component_clause, [], [f2817])).
fof(f2817, plain, (spl150_18 <=> (e15 = op1(e15, e13))), introduced(avatar_definition, [new_symbols(naming, [spl150_18])])).
fof(f25819, plain, ((e12 = op1(op1(e15, e13), e13)) | (~ spl150_9 | ~ spl150_89)), inference(forward_demodulation, [], [f24339, f2782])).
fof(f24339, plain, ((op1(e15, e14) = op1(op1(e15, e13), e13)) | ~ spl150_89), inference(forward_demodulation, [], [f601, f3115])).
fof(f601, plain, (op1(op1(e15, e13), e13) = op1(e15, op1(e13, e13))), inference(cnf_transformation, [], [f2])).
fof(f30487, plain, (~ spl150_161 | ~ spl150_89), inference(avatar_split_clause, [], [f24335, f3113, f3413])).
fof(f3413, plain, (spl150_161 <=> (e14 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl150_161])])).
fof(f24335, plain, (~ (e14 = op1(e11, e13)) | ~ spl150_89), inference(forward_demodulation, [], [f1097, f3115])).
fof(f1097, plain, ~ (op1(e11, e13) = op1(e13, e13)), inference(cnf_transformation, [], [f17])).
fof(f30450, plain, (~ spl150_133 | ~ spl150_127), inference(avatar_split_clause, [], [f26169, f3272, f3297])).
fof(f26169, plain, (~ (e10 = op1(e12, e11)) | ~ spl150_127), inference(forward_demodulation, [], [f1171, f3274])).
fof(f1171, plain, ~ (op1(e12, e11) = op1(e12, e12)), inference(cnf_transformation, [], [f17])).
fof(f30099, plain, (~ spl150_40 | ~ spl150_46), inference(avatar_split_clause, [], [f30098, f2934, f2909])).
fof(f2909, plain, (spl150_40 <=> (e13 = op1(e14, e15))), introduced(avatar_definition, [new_symbols(naming, [spl150_40])])).
fof(f30098, plain, (~ (e13 = op1(e14, e15)) | ~ spl150_46), inference(forward_demodulation, [], [f1210, f2936])).
fof(f1210, plain, ~ (op1(e14, e14) = op1(e14, e15)), inference(cnf_transformation, [], [f17])).
fof(f30054, plain, (~ spl150_41 | ~ spl150_71), inference(avatar_split_clause, [], [f30053, f3038, f2913])).
fof(f2913, plain, (spl150_41 <=> (e14 = op1(e14, e15))), introduced(avatar_definition, [new_symbols(naming, [spl150_41])])).
fof(f30053, plain, (~ (e14 = op1(e14, e15)) | ~ spl150_71), inference(forward_demodulation, [], [f1200, f3040])).
fof(f1200, plain, ~ (op1(e14, e10) = op1(e14, e15)), inference(cnf_transformation, [], [f17])).
fof(f30042, plain, (spl150_63 | ~ spl150_9 | ~ spl150_23 | ~ spl150_116 | ~ spl150_127 | ~ spl150_191), inference(avatar_split_clause, [], [f30041, f3538, f3272, f3226, f2838, f2780, f3005])).
fof(f30041, plain, ((e12 = op1(e14, e11)) | (~ spl150_9 | ~ spl150_23 | ~ spl150_116 | ~ spl150_127 | ~ spl150_191)), inference(forward_demodulation, [], [f30040, f2782])).
fof(f30040, plain, ((op1(e14, e11) = op1(e15, e14)) | (~ spl150_23 | ~ spl150_116 | ~ spl150_127 | ~ spl150_191)), inference(forward_demodulation, [], [f30039, f30003])).
fof(f30003, plain, ((op1(e14, e11) = op1(op1(e14, e12), e14)) | ~ spl150_116), inference(forward_demodulation, [], [f560, f3228])).
fof(f560, plain, (op1(op1(e14, e12), e14) = op1(e14, op1(e12, e14))), inference(cnf_transformation, [], [f2])).
fof(f30039, plain, ((op1(e15, e14) = op1(op1(e14, e12), e14)) | (~ spl150_23 | ~ spl150_127 | ~ spl150_191)), inference(forward_demodulation, [], [f28178, f26157])).
fof(f26157, plain, ((op1(e14, e12) = op1(e15, e10)) | (~ spl150_23 | ~ spl150_127)), inference(forward_demodulation, [], [f26156, f2840])).
fof(f26156, plain, ((op1(e15, e10) = op1(op1(e15, e12), e12)) | ~ spl150_127), inference(forward_demodulation, [], [f594, f3274])).
fof(f594, plain, (op1(op1(e15, e12), e12) = op1(e15, op1(e12, e12))), inference(cnf_transformation, [], [f2])).
fof(f28178, plain, ((op1(e15, e14) = op1(op1(e15, e10), e14)) | ~ spl150_191), inference(forward_demodulation, [], [f584, f3540])).
fof(f584, plain, (op1(op1(e15, e10), e14) = op1(e15, op1(e10, e14))), inference(cnf_transformation, [], [f2])).
fof(f29781, plain, (spl150_21 | ~ spl150_222), inference(avatar_contradiction_clause, [], [f29780])).
fof(f29780, plain, ($false | (spl150_21 | ~ spl150_222)), inference(subsumption_resolution, [], [f29723, f2831])).
fof(f2831, plain, (~ (e12 = op1(e15, e12)) | spl150_21), inference(avatar_component_clause, [], [f2830])).
fof(f2830, plain, (spl150_21 <=> (e12 = op1(e15, e12))), introduced(avatar_definition, [new_symbols(naming, [spl150_21])])).
fof(f29723, plain, ((e12 = op1(e15, e12)) | ~ spl150_222), inference(backward_demodulation, [], [f620, f3669])).
fof(f3669, plain, ((e15 = unit1) | ~ spl150_222), inference(avatar_component_clause, [], [f3667])).
fof(f3667, plain, (spl150_222 <=> (e15 = unit1)), introduced(avatar_definition, [new_symbols(naming, [spl150_222])])).
fof(f620, plain, (e12 = op1(unit1, e12)), inference(cnf_transformation, [], [f3])).
fof(f29668, plain, (~ spl150_292 | spl150_653), inference(avatar_contradiction_clause, [], [f29667])).
fof(f29667, plain, ($false | (~ spl150_292 | spl150_653)), inference(subsumption_resolution, [], [f29660, f6381])).
fof(f6381, plain, (~ (e23 = h5(e13)) | spl150_653), inference(avatar_component_clause, [], [f6379])).
fof(f6379, plain, (spl150_653 <=> (e23 = h5(e13))), introduced(avatar_definition, [new_symbols(naming, [spl150_653])])).
fof(f29660, plain, ((e23 = h5(e13)) | ~ spl150_292), inference(backward_demodulation, [], [f1510, f3961])).
fof(f3961, plain, ((e23 = op2(e25, e20)) | ~ spl150_292), inference(avatar_component_clause, [], [f3959])).
fof(f3959, plain, (spl150_292 <=> (e23 = op2(e25, e20))), introduced(avatar_definition, [new_symbols(naming, [spl150_292])])).
fof(f1510, plain, (op2(e25, e20) = h5(e13)), inference(cnf_transformation, [], [f30])).
fof(f30, plain, ((op2(op2(e25, e20), op2(e25, e20)) = h5(e14)) & (op2(e25, e20) = h5(e13)) & (op2(e25, op2(op2(e25, e20), op2(e25, e20))) = h5(e12)) & (h5(e10) = op2(op2(e25, op2(op2(e25, e20), op2(e25, e20))), op2(e25, op2(op2(e25, e20), op2(e25, e20))))) & (e25 = h5(e15)) & (e20 = h5(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG027+1.p', ax30)).
fof(f29657, plain, (~ spl150_1336 | ~ spl150_300), inference(avatar_split_clause, [], [f29649, f3992, f10449])).
fof(f3992, plain, (spl150_300 <=> (e25 = op2(e24, e25))), introduced(avatar_definition, [new_symbols(naming, [spl150_300])])).
fof(f29649, plain, (~ (e25 = h26(e13)) | ~ spl150_300), inference(backward_demodulation, [], [f5709, f3994])).
fof(f3994, plain, ((e25 = op2(e24, e25)) | ~ spl150_300), inference(avatar_component_clause, [], [f3992])).
fof(f5709, plain, ~ (op2(e24, e25) = h26(e13)), inference(backward_demodulation, [], [f1304, f1636])).
fof(f1304, plain, ~ (op2(e20, e25) = op2(e24, e25)), inference(cnf_transformation, [], [f18])).
fof(f29635, plain, (~ spl150_2052 | ~ spl150_324), inference(avatar_split_clause, [], [f29630, f4092, f15909])).
fof(f4092, plain, (spl150_324 <=> (e25 = op2(e24, e21))), introduced(avatar_definition, [new_symbols(naming, [spl150_324])])).
fof(f29630, plain, (~ (e25 = h8(e13)) | ~ spl150_324), inference(backward_demodulation, [], [f5282, f4094])).
fof(f4094, plain, ((e25 = op2(e24, e21)) | ~ spl150_324), inference(avatar_component_clause, [], [f4092])).
fof(f5282, plain, ~ (op2(e24, e21) = h8(e13)), inference(backward_demodulation, [], [f1253, f1528])).
fof(f1253, plain, ~ (op2(e23, e21) = op2(e24, e21)), inference(cnf_transformation, [], [f18])).
fof(f29529, plain, (spl150_375 | ~ spl150_479), inference(avatar_contradiction_clause, [], [f29528])).
fof(f29528, plain, ($false | (spl150_375 | ~ spl150_479)), inference(subsumption_resolution, [], [f29477, f4306])).
fof(f29477, plain, ((e22 = op2(e22, e24)) | ~ spl150_479), inference(backward_demodulation, [], [f904, f4740])).
fof(f4740, plain, ((e24 = unit2) | ~ spl150_479), inference(avatar_component_clause, [], [f4738])).
fof(f29453, plain, (~ spl150_347 | ~ spl150_508 | ~ spl150_954), inference(avatar_contradiction_clause, [], [f29452])).
fof(f29452, plain, ($false | (~ spl150_347 | ~ spl150_508 | ~ spl150_954)), inference(subsumption_resolution, [], [f29451, f1435])).
fof(f1435, plain, ~ (e24 = e25), inference(cnf_transformation, [], [f20])).
fof(f29451, plain, ((e24 = e25) | (~ spl150_347 | ~ spl150_508 | ~ spl150_954)), inference(backward_demodulation, [], [f910, f29437])).
fof(f29437, plain, ((e24 = op2(e25, unit2)) | (~ spl150_347 | ~ spl150_508 | ~ spl150_954)), inference(backward_demodulation, [], [f26781, f29421])).
fof(f29421, plain, ((unit2 = h17(e13)) | ~ spl150_508), inference(backward_demodulation, [], [f1582, f29416])).
fof(f29416, plain, ((op2(e21, e23) = unit2) | ~ spl150_508), inference(backward_demodulation, [], [f914, f4861])).
fof(f4861, plain, ((e23 = inv2(e21)) | ~ spl150_508), inference(avatar_component_clause, [], [f4859])).
fof(f4859, plain, (spl150_508 <=> (e23 = inv2(e21))), introduced(avatar_definition, [new_symbols(naming, [spl150_508])])).
fof(f26781, plain, ((e24 = op2(e25, h17(e13))) | (~ spl150_347 | ~ spl150_954)), inference(forward_demodulation, [], [f20079, f4190])).
fof(f20079, plain, ((op2(e23, e23) = op2(e25, h17(e13))) | ~ spl150_954), inference(backward_demodulation, [], [f5531, f7986])).
fof(f5531, plain, (op2(h10(e13), e23) = op2(e25, h17(e13))), inference(backward_demodulation, [], [f5327, f1582])).
fof(f5327, plain, (op2(e25, op2(e21, e23)) = op2(h10(e13), e23)), inference(backward_demodulation, [], [f872, f1540])).
fof(f872, plain, (op2(op2(e25, e21), e23) = op2(e25, op2(e21, e23))), inference(cnf_transformation, [], [f6])).
fof(f29410, plain, (spl150_920 | ~ spl150_424), inference(avatar_split_clause, [], [f29409, f4509, f7801])).
fof(f7801, plain, (spl150_920 <=> (e23 = h12(e13))), introduced(avatar_definition, [new_symbols(naming, [spl150_920])])).
fof(f29409, plain, ((e23 = h12(e13)) | ~ spl150_424), inference(forward_demodulation, [], [f1552, f4511])).
fof(f1552, plain, (op2(e21, e22) = h12(e13)), inference(cnf_transformation, [], [f37])).
fof(f37, plain, ((op2(op2(e21, e22), op2(e21, e22)) = h12(e14)) & (op2(e21, e22) = h12(e13)) & (op2(e21, op2(op2(e21, e22), op2(e21, e22))) = h12(e12)) & (h12(e10) = op2(op2(e21, op2(op2(e21, e22), op2(e21, e22))), op2(e21, op2(op2(e21, e22), op2(e21, e22))))) & (e21 = h12(e15)) & (e22 = h12(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG027+1.p', ax37)).
fof(f29262, plain, (~ spl150_385 | ~ spl150_483 | ~ spl150_656), inference(avatar_contradiction_clause, [], [f29261])).
fof(f29261, plain, ($false | (~ spl150_385 | ~ spl150_483 | ~ spl150_656)), inference(subsumption_resolution, [], [f29260, f1424])).
fof(f1424, plain, ~ (e20 = e24), inference(cnf_transformation, [], [f20])).
fof(f29260, plain, ((e20 = e24) | (~ spl150_385 | ~ spl150_483 | ~ spl150_656)), inference(backward_demodulation, [], [f907, f29255])).
fof(f29255, plain, ((e20 = op2(unit2, e24)) | (~ spl150_385 | ~ spl150_483 | ~ spl150_656)), inference(backward_demodulation, [], [f29184, f29240])).
fof(f29240, plain, ((unit2 = h28(e13)) | ~ spl150_483), inference(forward_demodulation, [], [f1648, f27367])).
fof(f27367, plain, ((op2(e22, e25) = unit2) | ~ spl150_483), inference(backward_demodulation, [], [f923, f4757])).
fof(f4757, plain, ((e22 = inv2(e25)) | ~ spl150_483), inference(avatar_component_clause, [], [f4755])).
fof(f4755, plain, (spl150_483 <=> (e22 = inv2(e25))), introduced(avatar_definition, [new_symbols(naming, [spl150_483])])).
fof(f923, plain, (unit2 = op2(inv2(e25), e25)), inference(cnf_transformation, [], [f8])).
fof(f907, plain, (e24 = op2(unit2, e24)), inference(cnf_transformation, [], [f7])).
fof(f29092, plain, (spl150_1495 | ~ spl150_414), inference(avatar_split_clause, [], [f29091, f4467, f11818])).
fof(f4467, plain, (spl150_414 <=> (e25 = op2(e21, e24))), introduced(avatar_definition, [new_symbols(naming, [spl150_414])])).
fof(f29091, plain, ((e25 = h22(e13)) | ~ spl150_414), inference(forward_demodulation, [], [f1612, f4469])).
fof(f4469, plain, ((e25 = op2(e21, e24)) | ~ spl150_414), inference(avatar_component_clause, [], [f4467])).
fof(f29066, plain, (spl150_292 | ~ spl150_316 | ~ spl150_385 | ~ spl150_850), inference(avatar_split_clause, [], [f29065, f7432, f4347, f4059, f3959])).
fof(f4059, plain, (spl150_316 <=> (e23 = op2(e24, e22))), introduced(avatar_definition, [new_symbols(naming, [spl150_316])])).
fof(f29065, plain, ((e23 = op2(e25, e20)) | (~ spl150_316 | ~ spl150_385 | ~ spl150_850)), inference(forward_demodulation, [], [f29064, f4061])).
fof(f4061, plain, ((e23 = op2(e24, e22)) | ~ spl150_316), inference(avatar_component_clause, [], [f4059])).
fof(f29064, plain, ((op2(e24, e22) = op2(e25, e20)) | (~ spl150_385 | ~ spl150_850)), inference(forward_demodulation, [], [f29063, f4349])).
fof(f29063, plain, ((op2(e24, e22) = op2(e25, op2(e22, e22))) | ~ spl150_850), inference(forward_demodulation, [], [f5471, f7433])).
fof(f5471, plain, (op2(e25, op2(e22, e22)) = op2(h15(e13), e22)), inference(backward_demodulation, [], [f877, f1570])).
fof(f877, plain, (op2(op2(e25, e22), e22) = op2(e25, op2(e22, e22))), inference(cnf_transformation, [], [f6])).
fof(f28968, plain, (~ spl150_199 | ~ spl150_127), inference(avatar_split_clause, [], [f26161, f3272, f3572])).
fof(f3572, plain, (spl150_199 <=> (e10 = op1(e10, e12))), introduced(avatar_definition, [new_symbols(naming, [spl150_199])])).
fof(f26161, plain, (~ (e10 = op1(e10, e12)) | ~ spl150_127), inference(forward_demodulation, [], [f1077, f3274])).
fof(f1077, plain, ~ (op1(e10, e12) = op1(e12, e12)), inference(cnf_transformation, [], [f17])).
fof(f28527, plain, (~ spl150_125 | ~ spl150_89), inference(avatar_split_clause, [], [f24333, f3113, f3263])).
fof(f3263, plain, (spl150_125 <=> (e14 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl150_125])])).
fof(f24333, plain, (~ (e14 = op1(e12, e13)) | ~ spl150_89), inference(forward_demodulation, [], [f1100, f3115])).
fof(f1100, plain, ~ (op1(e12, e13) = op1(e13, e13)), inference(cnf_transformation, [], [f17])).
fof(f27911, plain, (~ spl150_42 | ~ spl150_186), inference(avatar_split_clause, [], [f27895, f3517, f2917])).
fof(f2917, plain, (spl150_42 <=> (e15 = op1(e14, e15))), introduced(avatar_definition, [new_symbols(naming, [spl150_42])])).
fof(f27895, plain, (~ (e15 = op1(e14, e15)) | ~ spl150_186), inference(backward_demodulation, [], [f1124, f3519])).
fof(f1124, plain, ~ (op1(e10, e15) = op1(e14, e15)), inference(cnf_transformation, [], [f17])).
fof(f27817, plain, (~ spl150_9 | ~ spl150_221), inference(avatar_contradiction_clause, [], [f27816])).
fof(f27816, plain, ($false | (~ spl150_9 | ~ spl150_221)), inference(subsumption_resolution, [], [f27815, f1417])).
fof(f1417, plain, ~ (e12 = e15), inference(cnf_transformation, [], [f19])).
fof(f27815, plain, ((e12 = e15) | (~ spl150_9 | ~ spl150_221)), inference(forward_demodulation, [], [f27760, f2782])).
fof(f27760, plain, ((e15 = op1(e15, e14)) | ~ spl150_221), inference(backward_demodulation, [], [f627, f3665])).
fof(f3665, plain, ((e14 = unit1) | ~ spl150_221), inference(avatar_component_clause, [], [f3663])).
fof(f3663, plain, (spl150_221 <=> (e14 = unit1)), introduced(avatar_definition, [new_symbols(naming, [spl150_221])])).
fof(f627, plain, (e15 = op1(e15, unit1)), inference(cnf_transformation, [], [f3])).
fof(f27704, plain, (spl150_549 | ~ spl150_298), inference(avatar_split_clause, [], [f27698, f3984, f5826])).
fof(f5826, plain, (spl150_549 <=> (e23 = h30(e13))), introduced(avatar_definition, [new_symbols(naming, [spl150_549])])).
fof(f3984, plain, (spl150_298 <=> (e23 = op2(e24, e25))), introduced(avatar_definition, [new_symbols(naming, [spl150_298])])).
fof(f27698, plain, ((e23 = h30(e13)) | ~ spl150_298), inference(backward_demodulation, [], [f1660, f3986])).
fof(f3986, plain, ((e23 = op2(e24, e25)) | ~ spl150_298), inference(avatar_component_clause, [], [f3984])).
fof(f27667, plain, (spl150_609 | ~ spl150_407), inference(avatar_split_clause, [], [f27666, f4438, f6140])).
fof(f4438, plain, (spl150_407 <=> (e24 = op2(e21, e25))), introduced(avatar_definition, [new_symbols(naming, [spl150_407])])).
fof(f27666, plain, ((e24 = h27(e13)) | ~ spl150_407), inference(backward_demodulation, [], [f1642, f4440])).
fof(f4440, plain, ((e24 = op2(e21, e25)) | ~ spl150_407), inference(avatar_component_clause, [], [f4438])).
fof(f1642, plain, (op2(e21, e25) = h27(e13)), inference(cnf_transformation, [], [f52])).
fof(f52, plain, ((op2(op2(e21, e25), op2(e21, e25)) = h27(e14)) & (op2(e21, e25) = h27(e13)) & (op2(e21, op2(op2(e21, e25), op2(e21, e25))) = h27(e12)) & (h27(e10) = op2(op2(e21, op2(op2(e21, e25), op2(e21, e25))), op2(e21, op2(op2(e21, e25), op2(e21, e25))))) & (e21 = h27(e15)) & (e25 = h27(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG027+1.p', ax52)).
fof(f27437, plain, (~ spl150_478 | spl150_496), inference(avatar_contradiction_clause, [], [f27436])).
fof(f27436, plain, ($false | (~ spl150_478 | spl150_496)), inference(subsumption_resolution, [], [f27388, f4810])).
fof(f4810, plain, (~ (e23 = inv2(e23)) | spl150_496), inference(avatar_component_clause, [], [f4809])).
fof(f27388, plain, ((e23 = inv2(e23)) | ~ spl150_478), inference(backward_demodulation, [], [f988, f4736])).
fof(f4736, plain, ((e23 = unit2) | ~ spl150_478), inference(avatar_component_clause, [], [f4734])).
fof(f27125, plain, (~ spl150_385 | ~ spl150_434 | ~ spl150_507), inference(avatar_contradiction_clause, [], [f27124])).
fof(f27124, plain, ($false | (~ spl150_385 | ~ spl150_434 | ~ spl150_507)), inference(subsumption_resolution, [], [f27123, f1426])).
fof(f1426, plain, ~ (e21 = e22), inference(cnf_transformation, [], [f20])).
fof(f27123, plain, ((e21 = e22) | (~ spl150_385 | ~ spl150_434 | ~ spl150_507)), inference(forward_demodulation, [], [f27122, f4553])).
fof(f27122, plain, ((e22 = op2(e21, e20)) | (~ spl150_385 | ~ spl150_507)), inference(forward_demodulation, [], [f27121, f4349])).
fof(f27121, plain, ((e22 = op2(e21, op2(e22, e22))) | ~ spl150_507), inference(forward_demodulation, [], [f27120, f903])).
fof(f27120, plain, ((op2(e21, op2(e22, e22)) = op2(unit2, e22)) | ~ spl150_507), inference(forward_demodulation, [], [f5394, f24680])).
fof(f24680, plain, ((unit2 = h12(e13)) | ~ spl150_507), inference(backward_demodulation, [], [f1552, f24671])).
fof(f24671, plain, ((op2(e21, e22) = unit2) | ~ spl150_507), inference(backward_demodulation, [], [f914, f4857])).
fof(f4857, plain, ((e22 = inv2(e21)) | ~ spl150_507), inference(avatar_component_clause, [], [f4855])).
fof(f4855, plain, (spl150_507 <=> (e22 = inv2(e21))), introduced(avatar_definition, [new_symbols(naming, [spl150_507])])).
fof(f5394, plain, (op2(e21, op2(e22, e22)) = op2(h12(e13), e22)), inference(backward_demodulation, [], [f733, f1552])).
fof(f733, plain, (op2(op2(e21, e22), e22) = op2(e21, op2(e22, e22))), inference(cnf_transformation, [], [f6])).
fof(f26732, plain, (~ spl150_895 | ~ spl150_347), inference(avatar_split_clause, [], [f26731, f4188, f7669])).
fof(f7669, plain, (spl150_895 <=> (e24 = h13(e13))), introduced(avatar_definition, [new_symbols(naming, [spl150_895])])).
fof(f26731, plain, (~ (e24 = h13(e13)) | ~ spl150_347), inference(forward_demodulation, [], [f5432, f4190])).
fof(f5432, plain, ~ (op2(e23, e23) = h13(e13)), inference(backward_demodulation, [], [f1370, f1558])).
fof(f1558, plain, (op2(e23, e22) = h13(e13)), inference(cnf_transformation, [], [f38])).
fof(f38, plain, ((op2(op2(e23, e22), op2(e23, e22)) = h13(e14)) & (op2(e23, e22) = h13(e13)) & (op2(e23, op2(op2(e23, e22), op2(e23, e22))) = h13(e12)) & (h13(e10) = op2(op2(e23, op2(op2(e23, e22), op2(e23, e22))), op2(e23, op2(op2(e23, e22), op2(e23, e22))))) & (e23 = h13(e15)) & (e22 = h13(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG027+1.p', ax38)).
fof(f1370, plain, ~ (op2(e23, e22) = op2(e23, e23)), inference(cnf_transformation, [], [f18])).
fof(f26684, plain, (~ spl150_549 | ~ spl150_304), inference(avatar_split_clause, [], [f26683, f4009, f5826])).
fof(f26683, plain, (~ (e23 = h30(e13)) | ~ spl150_304), inference(forward_demodulation, [], [f5797, f4011])).
fof(f5797, plain, ~ (op2(e24, e24) = h30(e13)), inference(backward_demodulation, [], [f1390, f1660])).
fof(f1390, plain, ~ (op2(e24, e24) = op2(e24, e25)), inference(cnf_transformation, [], [f18])).
fof(f26646, plain, (~ spl150_322 | spl150_964), inference(avatar_contradiction_clause, [], [f26645])).
fof(f26645, plain, ($false | (~ spl150_322 | spl150_964)), inference(subsumption_resolution, [], [f26644, f8054])).
fof(f8054, plain, (~ (e23 = h9(e13)) | spl150_964), inference(avatar_component_clause, [], [f8052])).
fof(f8052, plain, (spl150_964 <=> (e23 = h9(e13))), introduced(avatar_definition, [new_symbols(naming, [spl150_964])])).
fof(f26644, plain, ((e23 = h9(e13)) | ~ spl150_322), inference(forward_demodulation, [], [f1534, f4086])).
fof(f4086, plain, ((e23 = op2(e24, e21)) | ~ spl150_322), inference(avatar_component_clause, [], [f4084])).
fof(f4084, plain, (spl150_322 <=> (e23 = op2(e24, e21))), introduced(avatar_definition, [new_symbols(naming, [spl150_322])])).
fof(f1534, plain, (op2(e24, e21) = h9(e13)), inference(cnf_transformation, [], [f34])).
fof(f34, plain, ((op2(op2(e24, e21), op2(e24, e21)) = h9(e14)) & (op2(e24, e21) = h9(e13)) & (op2(e24, op2(op2(e24, e21), op2(e24, e21))) = h9(e12)) & (h9(e10) = op2(op2(e24, op2(op2(e24, e21), op2(e24, e21))), op2(e24, op2(op2(e24, e21), op2(e24, e21))))) & (e24 = h9(e15)) & (e21 = h9(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG027+1.p', ax34)).
fof(f26102, plain, (~ spl150_124 | ~ spl150_196), inference(avatar_contradiction_clause, [], [f26101])).
fof(f26101, plain, ($false | (~ spl150_124 | ~ spl150_196)), inference(subsumption_resolution, [], [f26100, f3561])).
fof(f26100, plain, (~ (e13 = op1(e10, e13)) | ~ spl150_124), inference(forward_demodulation, [], [f1092, f3261])).
fof(f3261, plain, ((e13 = op1(e12, e13)) | ~ spl150_124), inference(avatar_component_clause, [], [f3259])).
fof(f3259, plain, (spl150_124 <=> (e13 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl150_124])])).
fof(f1092, plain, ~ (op1(e10, e13) = op1(e12, e13)), inference(cnf_transformation, [], [f17])).
fof(f26019, plain, (~ spl150_83 | ~ spl150_89), inference(avatar_split_clause, [], [f26018, f3113, f3088])).
fof(f26018, plain, (~ (e14 = op1(e13, e14)) | ~ spl150_89), inference(forward_demodulation, [], [f1193, f3115])).
fof(f1193, plain, ~ (op1(e13, e13) = op1(e13, e14)), inference(cnf_transformation, [], [f17])).
fof(f25823, plain, (~ spl150_17 | ~ spl150_89), inference(avatar_split_clause, [], [f25822, f3113, f2813])).
fof(f2813, plain, (spl150_17 <=> (e14 = op1(e15, e13))), introduced(avatar_definition, [new_symbols(naming, [spl150_17])])).
fof(f25822, plain, (~ (e14 = op1(e15, e13)) | ~ spl150_89), inference(forward_demodulation, [], [f1104, f3115])).
fof(f1104, plain, ~ (op1(e13, e13) = op1(e15, e13)), inference(cnf_transformation, [], [f17])).
fof(f25817, plain, (spl150_526 | ~ spl150_89), inference(avatar_split_clause, [], [f24353, f3113, f4993])).
fof(f24353, plain, ((op1(e13, e14) = op1(e14, e13)) | ~ spl150_89), inference(forward_demodulation, [], [f529, f3115])).
fof(f529, plain, (op1(op1(e13, e13), e13) = op1(e13, op1(e13, e13))), inference(cnf_transformation, [], [f2])).
fof(f25489, plain, (~ spl150_152 | ~ spl150_176), inference(avatar_split_clause, [], [f25453, f3476, f3376])).
fof(f3376, plain, (spl150_152 <=> (e11 = op1(e11, e14))), introduced(avatar_definition, [new_symbols(naming, [spl150_152])])).
fof(f25453, plain, (~ (e11 = op1(e11, e14)) | ~ spl150_176), inference(backward_demodulation, [], [f1154, f3478])).
fof(f1154, plain, ~ (op1(e11, e10) = op1(e11, e14)), inference(cnf_transformation, [], [f17])).
fof(f25224, plain, (~ spl150_920 | ~ spl150_406), inference(avatar_split_clause, [], [f25219, f4434, f7801])).
fof(f4434, plain, (spl150_406 <=> (e23 = op2(e21, e25))), introduced(avatar_definition, [new_symbols(naming, [spl150_406])])).
fof(f25219, plain, (~ (e23 = h12(e13)) | ~ spl150_406), inference(backward_demodulation, [], [f5405, f4436])).
fof(f4436, plain, ((e23 = op2(e21, e25)) | ~ spl150_406), inference(avatar_component_clause, [], [f4434])).
fof(f5405, plain, ~ (op2(e21, e25) = h12(e13)), inference(backward_demodulation, [], [f1342, f1552])).
fof(f1342, plain, ~ (op2(e21, e22) = op2(e21, e25)), inference(cnf_transformation, [], [f18])).
fof(f24730, plain, (spl150_895 | ~ spl150_498 | ~ spl150_656), inference(avatar_split_clause, [], [f24729, f6398, f4817, f7669])).
fof(f4817, plain, (spl150_498 <=> (e25 = inv2(e23))), introduced(avatar_definition, [new_symbols(naming, [spl150_498])])).
fof(f24729, plain, ((e24 = h13(e13)) | (~ spl150_498 | ~ spl150_656)), inference(forward_demodulation, [], [f24715, f907])).
fof(f24715, plain, ((op2(unit2, e24) = h13(e13)) | (~ spl150_498 | ~ spl150_656)), inference(backward_demodulation, [], [f20142, f24697])).
fof(f24697, plain, ((unit2 = h29(e13)) | ~ spl150_498), inference(backward_demodulation, [], [f1654, f24692])).
fof(f24692, plain, ((op2(e23, e25) = unit2) | ~ spl150_498), inference(backward_demodulation, [], [f918, f4819])).
fof(f4819, plain, ((e25 = inv2(e23)) | ~ spl150_498), inference(avatar_component_clause, [], [f4817])).
fof(f918, plain, (unit2 = op2(e23, inv2(e23))), inference(cnf_transformation, [], [f8])).
fof(f1654, plain, (op2(e23, e25) = h29(e13)), inference(cnf_transformation, [], [f54])).
fof(f54, plain, ((op2(op2(e23, e25), op2(e23, e25)) = h29(e14)) & (op2(e23, e25) = h29(e13)) & (op2(e23, op2(op2(e23, e25), op2(e23, e25))) = h29(e12)) & (h29(e10) = op2(op2(e23, op2(op2(e23, e25), op2(e23, e25))), op2(e23, op2(op2(e23, e25), op2(e23, e25))))) & (e23 = h29(e15)) & (e25 = h29(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG027+1.p', ax54)).
fof(f20142, plain, ((h13(e13) = op2(h29(e13), e24)) | ~ spl150_656), inference(forward_demodulation, [], [f20123, f1558])).
fof(f20123, plain, ((op2(e23, e22) = op2(h29(e13), e24)) | ~ spl150_656), inference(backward_demodulation, [], [f5781, f6399])).
fof(f5781, plain, (op2(e23, h10(e12)) = op2(h29(e13), e24)), inference(backward_demodulation, [], [f5343, f1654])).
fof(f5343, plain, (op2(op2(e23, e25), e24) = op2(e23, h10(e12))), inference(backward_demodulation, [], [f825, f5339])).
fof(f825, plain, (op2(op2(e23, e25), e24) = op2(e23, op2(e25, e24))), inference(cnf_transformation, [], [f6])).
fof(f24568, plain, (spl150_554 | ~ spl150_297), inference(avatar_split_clause, [], [f24567, f3980, f5852])).
fof(f3980, plain, (spl150_297 <=> (e22 = op2(e24, e25))), introduced(avatar_definition, [new_symbols(naming, [spl150_297])])).
fof(f24567, plain, ((e22 = h30(e13)) | ~ spl150_297), inference(forward_demodulation, [], [f1660, f3982])).
fof(f3982, plain, ((e22 = op2(e24, e25)) | ~ spl150_297), inference(avatar_component_clause, [], [f3980])).
fof(f24267, plain, (~ spl150_206 | ~ spl150_134), inference(avatar_split_clause, [], [f24266, f3301, f3601])).
fof(f3301, plain, (spl150_134 <=> (e11 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl150_134])])).
fof(f24266, plain, (~ (e11 = op1(e10, e11)) | ~ spl150_134), inference(forward_demodulation, [], [f1062, f3303])).
fof(f3303, plain, ((e11 = op1(e12, e11)) | ~ spl150_134), inference(avatar_component_clause, [], [f3301])).
fof(f1062, plain, ~ (op1(e10, e11) = op1(e12, e11)), inference(cnf_transformation, [], [f17])).
fof(f24209, plain, (spl150_199 | ~ spl150_219), inference(avatar_split_clause, [], [f22961, f3655, f3572])).
fof(f3655, plain, (spl150_219 <=> (e12 = unit1)), introduced(avatar_definition, [new_symbols(naming, [spl150_219])])).
fof(f22961, plain, ((e10 = op1(e10, e12)) | ~ spl150_219), inference(backward_demodulation, [], [f617, f3657])).
fof(f3657, plain, ((e12 = unit1) | ~ spl150_219), inference(avatar_component_clause, [], [f3655])).
fof(f24134, plain, (~ spl150_141 | ~ spl150_123), inference(avatar_split_clause, [], [f24133, f3255, f3330])).
fof(f3255, plain, (spl150_123 <=> (e12 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl150_123])])).
fof(f24133, plain, (~ (e12 = op1(e12, e10)) | ~ spl150_123), inference(forward_demodulation, [], [f1168, f3257])).
fof(f3257, plain, ((e12 = op1(e12, e13)) | ~ spl150_123), inference(avatar_component_clause, [], [f3255])).
fof(f1168, plain, ~ (op1(e12, e10) = op1(e12, e13)), inference(cnf_transformation, [], [f17])).
fof(f22933, plain, (spl150_850 | ~ spl150_281), inference(avatar_split_clause, [], [f22928, f3913, f7432])).
fof(f22928, plain, ((e24 = h15(e13)) | ~ spl150_281), inference(backward_demodulation, [], [f1570, f3915])).
fof(f22908, plain, (~ spl150_945 | ~ spl150_315), inference(avatar_split_clause, [], [f22904, f4055, f7933])).
fof(f7933, plain, (spl150_945 <=> (e22 = h11(e13))), introduced(avatar_definition, [new_symbols(naming, [spl150_945])])).
fof(f4055, plain, (spl150_315 <=> (e22 = op2(e24, e22))), introduced(avatar_definition, [new_symbols(naming, [spl150_315])])).
fof(f22904, plain, (~ (e22 = h11(e13)) | ~ spl150_315), inference(backward_demodulation, [], [f5374, f4057])).
fof(f4057, plain, ((e22 = op2(e24, e22)) | ~ spl150_315), inference(avatar_component_clause, [], [f4055])).
fof(f5374, plain, ~ (op2(e24, e22) = h11(e13)), inference(backward_demodulation, [], [f1259, f1546])).
fof(f1546, plain, (op2(e20, e22) = h11(e13)), inference(cnf_transformation, [], [f36])).
fof(f36, plain, ((op2(op2(e20, e22), op2(e20, e22)) = h11(e14)) & (op2(e20, e22) = h11(e13)) & (op2(e20, op2(op2(e20, e22), op2(e20, e22))) = h11(e12)) & (h11(e10) = op2(op2(e20, op2(op2(e20, e22), op2(e20, e22))), op2(e20, op2(op2(e20, e22), op2(e20, e22))))) & (e20 = h11(e15)) & (e22 = h11(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG027+1.p', ax36)).
fof(f1259, plain, ~ (op2(e20, e22) = op2(e24, e22)), inference(cnf_transformation, [], [f18])).
fof(f22897, plain, (spl150_969 | ~ spl150_321), inference(avatar_split_clause, [], [f22893, f4080, f8078])).
fof(f4080, plain, (spl150_321 <=> (e22 = op2(e24, e21))), introduced(avatar_definition, [new_symbols(naming, [spl150_321])])).
fof(f22893, plain, ((e22 = h9(e13)) | ~ spl150_321), inference(backward_demodulation, [], [f1534, f4082])).
fof(f4082, plain, ((e22 = op2(e24, e21)) | ~ spl150_321), inference(avatar_component_clause, [], [f4080])).
fof(f22885, plain, (spl150_2052 | ~ spl150_360), inference(avatar_split_clause, [], [f22882, f4242, f15909])).
fof(f22882, plain, ((e25 = h8(e13)) | ~ spl150_360), inference(backward_demodulation, [], [f1528, f4244])).
fof(f22850, plain, (~ spl150_820 | ~ spl150_405), inference(avatar_split_clause, [], [f22845, f4430, f7273])).
fof(f4430, plain, (spl150_405 <=> (e22 = op2(e21, e25))), introduced(avatar_definition, [new_symbols(naming, [spl150_405])])).
fof(f22845, plain, (~ (e22 = h17(e13)) | ~ spl150_405), inference(backward_demodulation, [], [f5523, f4432])).
fof(f4432, plain, ((e22 = op2(e21, e25)) | ~ spl150_405), inference(avatar_component_clause, [], [f4430])).
fof(f5523, plain, ~ (op2(e21, e25) = h17(e13)), inference(backward_demodulation, [], [f1344, f1582])).
fof(f1344, plain, ~ (op2(e21, e23) = op2(e21, e25)), inference(cnf_transformation, [], [f18])).
fof(f22842, plain, (spl150_710 | ~ spl150_412), inference(avatar_split_clause, [], [f22838, f4459, f6692])).
fof(f4459, plain, (spl150_412 <=> (e23 = op2(e21, e24))), introduced(avatar_definition, [new_symbols(naming, [spl150_412])])).
fof(f22838, plain, ((e23 = h22(e13)) | ~ spl150_412), inference(backward_demodulation, [], [f1612, f4461])).
fof(f4461, plain, ((e23 = op2(e21, e24)) | ~ spl150_412), inference(avatar_component_clause, [], [f4459])).
fof(f22833, plain, (spl150_820 | ~ spl150_417), inference(avatar_split_clause, [], [f22830, f4480, f7273])).
fof(f22830, plain, ((e22 = h17(e13)) | ~ spl150_417), inference(backward_demodulation, [], [f1582, f4482])).
fof(f22781, plain, (~ spl150_286 | ~ spl150_476), inference(avatar_contradiction_clause, [], [f22780])).
fof(f22780, plain, ($false | (~ spl150_286 | ~ spl150_476)), inference(subsumption_resolution, [], [f22779, f1434])).
fof(f1434, plain, ~ (e23 = e25), inference(cnf_transformation, [], [f20])).
fof(f22779, plain, ((e23 = e25) | (~ spl150_286 | ~ spl150_476)), inference(forward_demodulation, [], [f22721, f3936])).
fof(f22721, plain, ((e25 = op2(e25, e21)) | ~ spl150_476), inference(backward_demodulation, [], [f910, f4728])).
fof(f4728, plain, ((e21 = unit2) | ~ spl150_476), inference(avatar_component_clause, [], [f4726])).
fof(f4726, plain, (spl150_476 <=> (e21 = unit2)), introduced(avatar_definition, [new_symbols(naming, [spl150_476])])).
fof(f22710, plain, (spl150_477 | ~ spl150_267 | ~ spl150_485), inference(avatar_split_clause, [], [f22709, f4763, f3855, f4730])).
fof(f4763, plain, (spl150_485 <=> (e24 = inv2(e25))), introduced(avatar_definition, [new_symbols(naming, [spl150_485])])).
fof(f22709, plain, ((e22 = unit2) | (~ spl150_267 | ~ spl150_485)), inference(forward_demodulation, [], [f22705, f3857])).
fof(f22705, plain, ((op2(e25, e24) = unit2) | ~ spl150_485), inference(backward_demodulation, [], [f922, f4765])).
fof(f4765, plain, ((e24 = inv2(e25)) | ~ spl150_485), inference(avatar_component_clause, [], [f4763])).
fof(f22392, plain, (~ spl150_404 | ~ spl150_434), inference(avatar_split_clause, [], [f22373, f4551, f4426])).
fof(f4426, plain, (spl150_404 <=> (e21 = op2(e21, e25))), introduced(avatar_definition, [new_symbols(naming, [spl150_404])])).
fof(f22373, plain, (~ (e21 = op2(e21, e25)) | ~ spl150_434), inference(backward_demodulation, [], [f5107, f20178])).
fof(f20178, plain, ((e21 = h1(e13)) | ~ spl150_434), inference(backward_demodulation, [], [f1486, f4553])).
fof(f5107, plain, ~ (op2(e21, e25) = h1(e13)), inference(backward_demodulation, [], [f1335, f1486])).
fof(f1335, plain, ~ (op2(e21, e20) = op2(e21, e25)), inference(cnf_transformation, [], [f18])).
fof(f22340, plain, (spl150_385 | ~ spl150_656), inference(avatar_split_clause, [], [f20114, f6398, f4347])).
fof(f20114, plain, ((e20 = op2(e22, e22)) | ~ spl150_656), inference(backward_demodulation, [], [f5358, f6399])).
fof(f22333, plain, (spl150_347 | ~ spl150_954), inference(avatar_split_clause, [], [f20075, f7985, f4188])).
fof(f20075, plain, ((e24 = op2(e23, e23)) | ~ spl150_954), inference(backward_demodulation, [], [f5336, f7986])).
fof(f21775, plain, (~ spl150_136 | ~ spl150_28), inference(avatar_split_clause, [], [f21774, f2859, f3309])).
fof(f21774, plain, (~ (e13 = op1(e12, e11)) | ~ spl150_28), inference(forward_demodulation, [], [f1072, f2861])).
fof(f1072, plain, ~ (op1(e12, e11) = op1(e15, e11)), inference(cnf_transformation, [], [f17])).
fof(f21745, plain, (spl150_127 | ~ spl150_9), inference(avatar_split_clause, [], [f19881, f2780, f3272])).
fof(f19881, plain, ((e10 = op1(e12, e12)) | ~ spl150_9), inference(backward_demodulation, [], [f5081, f2782])).
fof(f5081, plain, (e10 = op1(op1(e15, e14), op1(e15, e14))), inference(forward_demodulation, [], [f1474, f1477])).
fof(f1477, plain, (e14 = op1(op1(e15, e11), op1(e15, e11))), inference(cnf_transformation, [], [f24])).
fof(f24, plain, ((e14 = op1(op1(e15, e11), op1(e15, e11))) & (e13 = op1(e15, e11)) & (e12 = op1(e15, op1(op1(e15, e11), op1(e15, e11)))) & (e10 = op1(op1(e15, op1(op1(e15, e11), op1(e15, e11))), op1(e15, op1(op1(e15, e11), op1(e15, e11)))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG027+1.p', ax24)).
fof(f1474, plain, (e10 = op1(op1(e15, op1(op1(e15, e11), op1(e15, e11))), op1(e15, op1(op1(e15, e11), op1(e15, e11))))), inference(cnf_transformation, [], [f24])).
fof(f21705, plain, (~ spl150_81 | ~ spl150_9), inference(avatar_split_clause, [], [f21704, f2780, f3080])).
fof(f21704, plain, (~ (e12 = op1(e13, e14)) | ~ spl150_9), inference(forward_demodulation, [], [f1119, f2782])).
fof(f1119, plain, ~ (op1(e13, e14) = op1(e15, e14)), inference(cnf_transformation, [], [f17])).
fof(f21030, plain, (~ spl150_158 | ~ spl150_176), inference(avatar_split_clause, [], [f21006, f3476, f3401])).
fof(f3401, plain, (spl150_158 <=> (e11 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl150_158])])).
fof(f21006, plain, (~ (e11 = op1(e11, e13)) | ~ spl150_176), inference(backward_demodulation, [], [f1153, f3478])).
fof(f1153, plain, ~ (op1(e11, e10) = op1(e11, e13)), inference(cnf_transformation, [], [f17])).
fof(f20777, plain, (~ spl150_28 | ~ spl150_218), inference(avatar_contradiction_clause, [], [f20776])).
fof(f20776, plain, ($false | (~ spl150_28 | ~ spl150_218)), inference(subsumption_resolution, [], [f20775, f1419])).
fof(f1419, plain, ~ (e13 = e15), inference(cnf_transformation, [], [f19])).
fof(f20775, plain, ((e13 = e15) | (~ spl150_28 | ~ spl150_218)), inference(forward_demodulation, [], [f20702, f2861])).
fof(f20702, plain, ((e15 = op1(e15, e11)) | ~ spl150_218), inference(backward_demodulation, [], [f627, f3653])).
fof(f3653, plain, ((e11 = unit1) | ~ spl150_218), inference(avatar_component_clause, [], [f3651])).
fof(f3651, plain, (spl150_218 <=> (e11 = unit1)), introduced(avatar_definition, [new_symbols(naming, [spl150_218])])).
fof(f20690, plain, (~ spl150_780 | ~ spl150_272), inference(avatar_split_clause, [], [f20684, f3876, f7062])).
fof(f20684, plain, (~ (e21 = h19(e13)) | ~ spl150_272), inference(backward_demodulation, [], [f5571, f3878])).
fof(f5571, plain, ~ (op2(e25, e23) = h19(e13)), inference(backward_demodulation, [], [f1285, f1594])).
fof(f1285, plain, ~ (op2(e24, e23) = op2(e25, e23)), inference(cnf_transformation, [], [f18])).
fof(f20685, plain, (spl150_760 | ~ spl150_272), inference(avatar_split_clause, [], [f20679, f3876, f6956])).
fof(f20679, plain, ((e21 = h20(e13)) | ~ spl150_272), inference(backward_demodulation, [], [f1600, f3878])).
fof(f20669, plain, (spl150_1376 | ~ spl150_294), inference(avatar_split_clause, [], [f20667, f3967, f10948])).
fof(f3967, plain, (spl150_294 <=> (e25 = op2(e25, e20))), introduced(avatar_definition, [new_symbols(naming, [spl150_294])])).
fof(f20667, plain, ((e25 = h5(e13)) | ~ spl150_294), inference(backward_demodulation, [], [f1510, f3969])).
fof(f3969, plain, ((e25 = op2(e25, e20)) | ~ spl150_294), inference(avatar_component_clause, [], [f3967])).
fof(f20663, plain, (~ spl150_885 | ~ spl150_296), inference(avatar_split_clause, [], [f20657, f3976, f7617])).
fof(f7617, plain, (spl150_885 <=> (e21 = h14(e13))), introduced(avatar_definition, [new_symbols(naming, [spl150_885])])).
fof(f3976, plain, (spl150_296 <=> (e21 = op2(e24, e25))), introduced(avatar_definition, [new_symbols(naming, [spl150_296])])).
fof(f20657, plain, (~ (e21 = h14(e13)) | ~ spl150_296), inference(backward_demodulation, [], [f5456, f3978])).
fof(f3978, plain, ((e21 = op2(e24, e25)) | ~ spl150_296), inference(avatar_component_clause, [], [f3976])).
fof(f5456, plain, ~ (op2(e24, e25) = h14(e13)), inference(backward_demodulation, [], [f1387, f1564])).
fof(f1387, plain, ~ (op2(e24, e22) = op2(e24, e25)), inference(cnf_transformation, [], [f18])).
fof(f20661, plain, (spl150_559 | ~ spl150_296), inference(avatar_split_clause, [], [f20655, f3976, f5878])).
fof(f20655, plain, ((e21 = h30(e13)) | ~ spl150_296), inference(backward_demodulation, [], [f1660, f3978])).
fof(f20642, plain, (spl150_885 | ~ spl150_314), inference(avatar_split_clause, [], [f20639, f4051, f7617])).
fof(f4051, plain, (spl150_314 <=> (e21 = op2(e24, e22))), introduced(avatar_definition, [new_symbols(naming, [spl150_314])])).
fof(f20639, plain, ((e21 = h14(e13)) | ~ spl150_314), inference(backward_demodulation, [], [f1564, f4053])).
fof(f4053, plain, ((e21 = op2(e24, e22)) | ~ spl150_314), inference(avatar_component_clause, [], [f4051])).
fof(f20636, plain, (~ spl150_1039 | ~ spl150_320), inference(avatar_split_clause, [], [f20632, f4076, f8447])).
fof(f4076, plain, (spl150_320 <=> (e21 = op2(e24, e21))), introduced(avatar_definition, [new_symbols(naming, [spl150_320])])).
fof(f20632, plain, (~ (e21 = h6(e13)) | ~ spl150_320), inference(backward_demodulation, [], [f5225, f4078])).
fof(f4078, plain, ((e21 = op2(e24, e21)) | ~ spl150_320), inference(avatar_component_clause, [], [f4076])).
fof(f5225, plain, ~ (op2(e24, e21) = h6(e13)), inference(backward_demodulation, [], [f1244, f1516])).
fof(f1244, plain, ~ (op2(e20, e21) = op2(e24, e21)), inference(cnf_transformation, [], [f18])).
fof(f20617, plain, (~ spl150_299 | ~ spl150_329), inference(avatar_split_clause, [], [f20599, f4113, f3988])).
fof(f3988, plain, (spl150_299 <=> (e24 = op2(e24, e25))), introduced(avatar_definition, [new_symbols(naming, [spl150_299])])).
fof(f20599, plain, (~ (e24 = op2(e24, e25)) | ~ spl150_329), inference(backward_demodulation, [], [f5185, f20593])).
fof(f20593, plain, ((e24 = h4(e13)) | ~ spl150_329), inference(backward_demodulation, [], [f1504, f4115])).
fof(f5185, plain, ~ (op2(e24, e25) = h4(e13)), inference(backward_demodulation, [], [f1380, f1504])).
fof(f1380, plain, ~ (op2(e24, e20) = op2(e24, e25)), inference(cnf_transformation, [], [f18])).
fof(f20615, plain, (~ spl150_317 | ~ spl150_329), inference(avatar_split_clause, [], [f20597, f4113, f4063])).
fof(f4063, plain, (spl150_317 <=> (e24 = op2(e24, e22))), introduced(avatar_definition, [new_symbols(naming, [spl150_317])])).
fof(f20597, plain, (~ (e24 = op2(e24, e22)) | ~ spl150_329), inference(backward_demodulation, [], [f5182, f20593])).
fof(f5182, plain, ~ (op2(e24, e22) = h4(e13)), inference(backward_demodulation, [], [f1377, f1504])).
fof(f1377, plain, ~ (op2(e24, e20) = op2(e24, e22)), inference(cnf_transformation, [], [f18])).
fof(f20614, plain, (~ spl150_323 | ~ spl150_329), inference(avatar_split_clause, [], [f20596, f4113, f4088])).
fof(f4088, plain, (spl150_323 <=> (e24 = op2(e24, e21))), introduced(avatar_definition, [new_symbols(naming, [spl150_323])])).
fof(f20596, plain, (~ (e24 = op2(e24, e21)) | ~ spl150_329), inference(backward_demodulation, [], [f5181, f20593])).
fof(f5181, plain, ~ (op2(e24, e21) = h4(e13)), inference(backward_demodulation, [], [f1376, f1504])).
fof(f1376, plain, ~ (op2(e24, e20) = op2(e24, e21)), inference(cnf_transformation, [], [f18])).
fof(f20363, plain, (spl150_340 | ~ spl150_286 | ~ spl150_410 | ~ spl150_954), inference(avatar_split_clause, [], [f20348, f7985, f4451, f3934, f4159])).
fof(f4451, plain, (spl150_410 <=> (e21 = op2(e21, e24))), introduced(avatar_definition, [new_symbols(naming, [spl150_410])])).
fof(f20348, plain, ((e23 = op2(e23, e24)) | (~ spl150_286 | ~ spl150_410 | ~ spl150_954)), inference(backward_demodulation, [], [f1624, f20347])).
fof(f20347, plain, ((e23 = h24(e13)) | (~ spl150_286 | ~ spl150_410 | ~ spl150_954)), inference(forward_demodulation, [], [f20327, f3936])).
fof(f20327, plain, ((op2(e25, e21) = h24(e13)) | (~ spl150_410 | ~ spl150_954)), inference(backward_demodulation, [], [f20099, f20309])).
fof(f20309, plain, ((e21 = h22(e13)) | ~ spl150_410), inference(backward_demodulation, [], [f1612, f4453])).
fof(f4453, plain, ((e21 = op2(e21, e24)) | ~ spl150_410), inference(avatar_component_clause, [], [f4451])).
fof(f20173, plain, (spl150_1336 | ~ spl150_444), inference(avatar_split_clause, [], [f20168, f4592, f10449])).
fof(f4592, plain, (spl150_444 <=> (e25 = op2(e20, e25))), introduced(avatar_definition, [new_symbols(naming, [spl150_444])])).
fof(f20168, plain, ((e25 = h26(e13)) | ~ spl150_444), inference(backward_demodulation, [], [f1636, f4594])).
fof(f4594, plain, ((e25 = op2(e20, e25)) | ~ spl150_444), inference(avatar_component_clause, [], [f4592])).
fof(f20164, plain, (spl150_725 | ~ spl150_449), inference(avatar_split_clause, [], [f20160, f4613, f6771])).
fof(f4613, plain, (spl150_449 <=> (e24 = op2(e20, e24))), introduced(avatar_definition, [new_symbols(naming, [spl150_449])])).
fof(f20160, plain, ((e24 = h21(e13)) | ~ spl150_449), inference(backward_demodulation, [], [f1606, f4615])).
fof(f4615, plain, ((e24 = op2(e20, e24)) | ~ spl150_449), inference(avatar_component_clause, [], [f4613])).
fof(f20153, plain, (spl150_945 | ~ spl150_459), inference(avatar_split_clause, [], [f20151, f4655, f7933])).
fof(f20151, plain, ((e22 = h11(e13)) | ~ spl150_459), inference(backward_demodulation, [], [f1546, f4657])).
fof(f20150, plain, (spl150_1039 | ~ spl150_464), inference(avatar_split_clause, [], [f20149, f4676, f8447])).
fof(f20149, plain, ((e21 = h6(e13)) | ~ spl150_464), inference(backward_demodulation, [], [f1516, f4678])).
fof(f4678, plain, ((e21 = op2(e20, e21)) | ~ spl150_464), inference(avatar_component_clause, [], [f4676])).
fof(f20146, plain, (spl150_281 | ~ spl150_486 | ~ spl150_656), inference(avatar_split_clause, [], [f20127, f6398, f4767, f3913])).
fof(f4767, plain, (spl150_486 <=> (e25 = inv2(e25))), introduced(avatar_definition, [new_symbols(naming, [spl150_486])])).
fof(f20127, plain, ((e24 = op2(e25, e22)) | (~ spl150_486 | ~ spl150_656)), inference(backward_demodulation, [], [f18235, f6399])).
fof(f18235, plain, ((e24 = op2(e25, h10(e12))) | ~ spl150_486), inference(forward_demodulation, [], [f18215, f907])).
fof(f18215, plain, ((op2(unit2, e24) = op2(e25, h10(e12))) | ~ spl150_486), inference(backward_demodulation, [], [f5349, f18207])).
fof(f18207, plain, ((op2(e25, e25) = unit2) | ~ spl150_486), inference(backward_demodulation, [], [f922, f4769])).
fof(f4769, plain, ((e25 = inv2(e25)) | ~ spl150_486), inference(avatar_component_clause, [], [f4767])).
fof(f5349, plain, (op2(op2(e25, e25), e24) = op2(e25, h10(e12))), inference(backward_demodulation, [], [f897, f5339])).
fof(f897, plain, (op2(op2(e25, e25), e24) = op2(e25, op2(e25, e24))), inference(cnf_transformation, [], [f6])).
fof(f20133, plain, (~ spl150_279 | ~ spl150_656), inference(avatar_split_clause, [], [f20112, f6398, f3905])).
fof(f20112, plain, (~ (e22 = op2(e25, e22)) | ~ spl150_656), inference(backward_demodulation, [], [f5355, f6399])).
fof(f5355, plain, ~ (op2(e25, e22) = h10(e12)), inference(backward_demodulation, [], [f1401, f5339])).
fof(f1401, plain, ~ (op2(e25, e22) = op2(e25, e24)), inference(cnf_transformation, [], [f18])).
fof(f20131, plain, (~ spl150_375 | ~ spl150_656), inference(avatar_split_clause, [], [f20110, f6398, f4305])).
fof(f20110, plain, (~ (e22 = op2(e22, e24)) | ~ spl150_656), inference(backward_demodulation, [], [f5352, f6399])).
fof(f5352, plain, ~ (op2(e22, e24) = h10(e12)), inference(backward_demodulation, [], [f1297, f5339])).
fof(f1297, plain, ~ (op2(e22, e24) = op2(e25, e24)), inference(cnf_transformation, [], [f18])).
fof(f20130, plain, (~ spl150_411 | ~ spl150_656), inference(avatar_split_clause, [], [f20109, f6398, f4455])).
fof(f4455, plain, (spl150_411 <=> (e22 = op2(e21, e24))), introduced(avatar_definition, [new_symbols(naming, [spl150_411])])).
fof(f20109, plain, (~ (e22 = op2(e21, e24)) | ~ spl150_656), inference(backward_demodulation, [], [f5351, f6399])).
fof(f5351, plain, ~ (op2(e21, e24) = h10(e12)), inference(backward_demodulation, [], [f1294, f5339])).
fof(f1294, plain, ~ (op2(e21, e24) = op2(e25, e24)), inference(cnf_transformation, [], [f18])).
fof(f20107, plain, (spl150_272 | ~ spl150_486 | ~ spl150_954), inference(avatar_split_clause, [], [f20089, f7985, f4767, f3876])).
fof(f20089, plain, ((e21 = op2(e25, e23)) | (~ spl150_486 | ~ spl150_954)), inference(backward_demodulation, [], [f18234, f7986])).
fof(f18234, plain, ((e21 = op2(e25, h10(e13))) | ~ spl150_486), inference(forward_demodulation, [], [f18213, f901])).
fof(f18213, plain, ((op2(unit2, e21) = op2(e25, h10(e13))) | ~ spl150_486), inference(backward_demodulation, [], [f5330, f18207])).
fof(f5330, plain, (op2(op2(e25, e25), e21) = op2(e25, h10(e13))), inference(backward_demodulation, [], [f894, f1540])).
fof(f894, plain, (op2(op2(e25, e25), e21) = op2(e25, op2(e25, e21))), inference(cnf_transformation, [], [f6])).
fof(f20106, plain, (spl150_360 | ~ spl150_506 | ~ spl150_954), inference(avatar_split_clause, [], [f20088, f7985, f4851, f4242])).
fof(f20088, plain, ((e25 = op2(e23, e21)) | (~ spl150_506 | ~ spl150_954)), inference(backward_demodulation, [], [f18075, f7986])).
fof(f18075, plain, ((e25 = op2(h10(e13), e21)) | ~ spl150_506), inference(forward_demodulation, [], [f18059, f910])).
fof(f18059, plain, ((op2(e25, unit2) = op2(h10(e13), e21)) | ~ spl150_506), inference(backward_demodulation, [], [f5325, f18037])).
fof(f18037, plain, ((op2(e21, e21) = unit2) | ~ spl150_506), inference(backward_demodulation, [], [f914, f4853])).
fof(f5325, plain, (op2(e25, op2(e21, e21)) = op2(h10(e13), e21)), inference(backward_demodulation, [], [f870, f1540])).
fof(f870, plain, (op2(op2(e25, e21), e21) = op2(e25, op2(e21, e21))), inference(cnf_transformation, [], [f6])).
fof(f19997, plain, (~ spl150_157 | ~ spl150_49), inference(avatar_split_clause, [], [f19996, f2947, f3397])).
fof(f3397, plain, (spl150_157 <=> (e10 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl150_157])])).
fof(f19996, plain, (~ (e10 = op1(e11, e13)) | ~ spl150_49), inference(forward_demodulation, [], [f1098, f2949])).
fof(f1098, plain, ~ (op1(e11, e13) = op1(e14, e13)), inference(cnf_transformation, [], [f17])).
fof(f19991, plain, (~ spl150_151 | ~ spl150_79), inference(avatar_split_clause, [], [f19990, f3072, f3372])).
fof(f3372, plain, (spl150_151 <=> (e10 = op1(e11, e14))), introduced(avatar_definition, [new_symbols(naming, [spl150_151])])).
fof(f19990, plain, (~ (e10 = op1(e11, e14)) | ~ spl150_79), inference(forward_demodulation, [], [f1112, f3074])).
fof(f1112, plain, ~ (op1(e11, e14) = op1(e13, e14)), inference(cnf_transformation, [], [f17])).
fof(f19989, plain, (~ spl150_153 | ~ spl150_9), inference(avatar_split_clause, [], [f19988, f2780, f3380])).
fof(f3380, plain, (spl150_153 <=> (e12 = op1(e11, e14))), introduced(avatar_definition, [new_symbols(naming, [spl150_153])])).
fof(f19988, plain, (~ (e12 = op1(e11, e14)) | ~ spl150_9), inference(forward_demodulation, [], [f1114, f2782])).
fof(f1114, plain, ~ (op1(e11, e14) = op1(e15, e14)), inference(cnf_transformation, [], [f17])).
fof(f19951, plain, (~ spl150_97 | ~ spl150_79), inference(avatar_split_clause, [], [f19950, f3072, f3147])).
fof(f19950, plain, (~ (e10 = op1(e13, e11)) | ~ spl150_79), inference(forward_demodulation, [], [f1188, f3074])).
fof(f1188, plain, ~ (op1(e13, e11) = op1(e13, e14)), inference(cnf_transformation, [], [f17])).
fof(f19945, plain, (spl150_89 | ~ spl150_28), inference(avatar_split_clause, [], [f19816, f2859, f3113])).
fof(f19816, plain, ((e14 = op1(e13, e13)) | ~ spl150_28), inference(backward_demodulation, [], [f1477, f2861])).
fof(f19918, plain, (~ spl150_34 | ~ spl150_28), inference(avatar_split_clause, [], [f19917, f2859, f2884])).
fof(f19917, plain, (~ (e13 = op1(e15, e10)) | ~ spl150_28), inference(forward_demodulation, [], [f1211, f2861])).
fof(f1211, plain, ~ (op1(e15, e10) = op1(e15, e11)), inference(cnf_transformation, [], [f17])).
fof(f19904, plain, (~ spl150_21 | ~ spl150_9), inference(avatar_split_clause, [], [f19903, f2780, f2830])).
fof(f19903, plain, (~ (e12 = op1(e15, e12)) | ~ spl150_9), inference(forward_demodulation, [], [f1221, f2782])).
fof(f1221, plain, ~ (op1(e15, e12) = op1(e15, e14)), inference(cnf_transformation, [], [f17])).
fof(f19901, plain, (~ spl150_15 | ~ spl150_9), inference(avatar_split_clause, [], [f19900, f2780, f2805])).
fof(f19900, plain, (~ (e12 = op1(e15, e13)) | ~ spl150_9), inference(forward_demodulation, [], [f1223, f2782])).
fof(f1223, plain, ~ (op1(e15, e13) = op1(e15, e14)), inference(cnf_transformation, [], [f17])).
fof(f19830, plain, (~ spl150_16 | ~ spl150_28), inference(avatar_split_clause, [], [f19814, f2859, f2809])).
fof(f2809, plain, (spl150_16 <=> (e13 = op1(e15, e13))), introduced(avatar_definition, [new_symbols(naming, [spl150_16])])).
fof(f19814, plain, (~ (e13 = op1(e15, e13)) | ~ spl150_28), inference(backward_demodulation, [], [f1217, f2861])).
fof(f1217, plain, ~ (op1(e15, e11) = op1(e15, e13)), inference(cnf_transformation, [], [f17])).
fof(f19739, plain, (~ spl150_13 | ~ spl150_49), inference(avatar_split_clause, [], [f19726, f2947, f2797])).
fof(f2797, plain, (spl150_13 <=> (e10 = op1(e15, e13))), introduced(avatar_definition, [new_symbols(naming, [spl150_13])])).
fof(f19726, plain, (~ (e10 = op1(e15, e13)) | ~ spl150_49), inference(backward_demodulation, [], [f1105, f2949])).
fof(f1105, plain, ~ (op1(e14, e13) = op1(e15, e13)), inference(cnf_transformation, [], [f17])).
fof(f18989, plain, (spl150_169 | ~ spl150_217 | ~ spl150_248), inference(avatar_split_clause, [], [f18940, f3776, f3647, f3447])).
fof(f3776, plain, (spl150_248 <=> (e11 = inv1(e11))), introduced(avatar_definition, [new_symbols(naming, [spl150_248])])).
fof(f18940, plain, ((e10 = op1(e11, e11)) | (~ spl150_217 | ~ spl150_248)), inference(backward_demodulation, [], [f18686, f3649])).
fof(f18686, plain, ((op1(e11, e11) = unit1) | ~ spl150_248), inference(backward_demodulation, [], [f631, f3778])).
fof(f3778, plain, ((e11 = inv1(e11)) | ~ spl150_248), inference(avatar_component_clause, [], [f3776])).
fof(f631, plain, (unit1 = op1(e11, inv1(e11))), inference(cnf_transformation, [], [f4])).
fof(f18977, plain, (spl150_36 | ~ spl150_217), inference(avatar_split_clause, [], [f18903, f3647, f2892])).
fof(f18903, plain, ((e15 = op1(e15, e10)) | ~ spl150_217), inference(backward_demodulation, [], [f627, f3649])).
fof(f18976, plain, (spl150_186 | ~ spl150_217), inference(avatar_split_clause, [], [f18902, f3647, f3517])).
fof(f18902, plain, ((e15 = op1(e10, e15)) | ~ spl150_217), inference(backward_demodulation, [], [f626, f3649])).
fof(f626, plain, (e15 = op1(unit1, e15)), inference(cnf_transformation, [], [f3])).
fof(f18975, plain, (spl150_71 | ~ spl150_217), inference(avatar_split_clause, [], [f18901, f3647, f3038])).
fof(f18901, plain, ((e14 = op1(e14, e10)) | ~ spl150_217), inference(backward_demodulation, [], [f625, f3649])).
fof(f625, plain, (e14 = op1(e14, unit1)), inference(cnf_transformation, [], [f3])).
fof(f18974, plain, (spl150_191 | ~ spl150_217), inference(avatar_split_clause, [], [f18900, f3647, f3538])).
fof(f18900, plain, ((e14 = op1(e10, e14)) | ~ spl150_217), inference(backward_demodulation, [], [f624, f3649])).
fof(f18973, plain, (spl150_106 | ~ spl150_217), inference(avatar_split_clause, [], [f18899, f3647, f3184])).
fof(f18899, plain, ((e13 = op1(e13, e10)) | ~ spl150_217), inference(backward_demodulation, [], [f623, f3649])).
fof(f623, plain, (e13 = op1(e13, unit1)), inference(cnf_transformation, [], [f3])).
fof(f18972, plain, (spl150_196 | ~ spl150_217), inference(avatar_split_clause, [], [f18898, f3647, f3559])).
fof(f18898, plain, ((e13 = op1(e10, e13)) | ~ spl150_217), inference(backward_demodulation, [], [f622, f3649])).
fof(f622, plain, (e13 = op1(unit1, e13)), inference(cnf_transformation, [], [f3])).
fof(f18971, plain, (spl150_141 | ~ spl150_217), inference(avatar_split_clause, [], [f18897, f3647, f3330])).
fof(f18897, plain, ((e12 = op1(e12, e10)) | ~ spl150_217), inference(backward_demodulation, [], [f621, f3649])).
fof(f621, plain, (e12 = op1(e12, unit1)), inference(cnf_transformation, [], [f3])).
fof(f18970, plain, (spl150_201 | ~ spl150_217), inference(avatar_split_clause, [], [f18896, f3647, f3580])).
fof(f18896, plain, ((e12 = op1(e10, e12)) | ~ spl150_217), inference(backward_demodulation, [], [f620, f3649])).
fof(f18969, plain, (spl150_176 | ~ spl150_217), inference(avatar_split_clause, [], [f18895, f3647, f3476])).
fof(f18895, plain, ((e11 = op1(e11, e10)) | ~ spl150_217), inference(backward_demodulation, [], [f619, f3649])).
fof(f619, plain, (e11 = op1(e11, unit1)), inference(cnf_transformation, [], [f3])).
fof(f18968, plain, (spl150_206 | ~ spl150_217), inference(avatar_split_clause, [], [f18894, f3647, f3601])).
fof(f18894, plain, ((e11 = op1(e10, e11)) | ~ spl150_217), inference(backward_demodulation, [], [f618, f3649])).
fof(f618, plain, (e11 = op1(unit1, e11)), inference(cnf_transformation, [], [f3])).
fof(f18681, plain, (~ spl150_247 | ~ spl150_253), inference(avatar_split_clause, [], [f18615, f3797, f3772])).
fof(f3772, plain, (spl150_247 <=> (e10 = inv1(e11))), introduced(avatar_definition, [new_symbols(naming, [spl150_247])])).
fof(f18615, plain, (~ (e10 = inv1(e11)) | ~ spl150_253), inference(backward_demodulation, [], [f973, f3799])).
fof(f3799, plain, ((e10 = inv1(e10)) | ~ spl150_253), inference(avatar_component_clause, [], [f3797])).
fof(f973, plain, ~ (inv1(e10) = inv1(e11)), inference(cnf_transformation, [], [f12])).
fof(f12, plain, (~ (inv1(e14) = inv1(e15)) & ~ (inv1(e13) = inv1(e15)) & ~ (inv1(e13) = inv1(e14)) & ~ (inv1(e12) = inv1(e15)) & ~ (inv1(e12) = inv1(e14)) & ~ (inv1(e12) = inv1(e13)) & ~ (inv1(e11) = inv1(e15)) & ~ (inv1(e11) = inv1(e14)) & ~ (inv1(e11) = inv1(e13)) & ~ (inv1(e11) = inv1(e12)) & ~ (inv1(e10) = inv1(e15)) & ~ (inv1(e10) = inv1(e14)) & ~ (inv1(e10) = inv1(e13)) & ~ (inv1(e10) = inv1(e12)) & ~ (inv1(e10) = inv1(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG027+1.p', ax12)).
fof(f18611, plain, (spl150_656 | ~ spl150_267), inference(avatar_split_clause, [], [f18609, f3855, f6398])).
fof(f18609, plain, ((e22 = h10(e12)) | ~ spl150_267), inference(backward_demodulation, [], [f5339, f3857])).
fof(f18575, plain, (~ spl150_964 | ~ spl150_286), inference(avatar_split_clause, [], [f18570, f3934, f8052])).
fof(f18570, plain, (~ (e23 = h9(e13)) | ~ spl150_286), inference(backward_demodulation, [], [f5307, f3936])).
fof(f5307, plain, ~ (op2(e25, e21) = h9(e13)), inference(backward_demodulation, [], [f1255, f1534])).
fof(f1255, plain, ~ (op2(e24, e21) = op2(e25, e21)), inference(cnf_transformation, [], [f18])).
fof(f18572, plain, (~ spl150_653 | ~ spl150_286), inference(avatar_split_clause, [], [f18566, f3934, f6379])).
fof(f18566, plain, (~ (e23 = h5(e13)) | ~ spl150_286), inference(backward_demodulation, [], [f5205, f3936])).
fof(f5205, plain, ~ (op2(e25, e21) = h5(e13)), inference(backward_demodulation, [], [f1391, f1510])).
fof(f1391, plain, ~ (op2(e25, e20) = op2(e25, e21)), inference(cnf_transformation, [], [f18])).
fof(f18571, plain, (spl150_954 | ~ spl150_286), inference(avatar_split_clause, [], [f18565, f3934, f7985])).
fof(f18565, plain, ((e23 = h10(e13)) | ~ spl150_286), inference(backward_demodulation, [], [f1540, f3936])).
fof(f18551, plain, (~ spl150_785 | ~ spl150_295), inference(avatar_split_clause, [], [f18542, f3972, f7088])).
fof(f3972, plain, (spl150_295 <=> (e20 = op2(e24, e25))), introduced(avatar_definition, [new_symbols(naming, [spl150_295])])).
fof(f18542, plain, (~ (e20 = h19(e13)) | ~ spl150_295), inference(backward_demodulation, [], [f5573, f3974])).
fof(f3974, plain, ((e20 = op2(e24, e25)) | ~ spl150_295), inference(avatar_component_clause, [], [f3972])).
fof(f5573, plain, ~ (op2(e24, e25) = h19(e13)), inference(backward_demodulation, [], [f1389, f1594])).
fof(f1389, plain, ~ (op2(e24, e23) = op2(e24, e25)), inference(cnf_transformation, [], [f18])).
fof(f18534, plain, (~ spl150_979 | ~ spl150_307), inference(avatar_split_clause, [], [f18527, f4022, f8131])).
fof(f8131, plain, (spl150_979 <=> (e20 = h9(e13))), introduced(avatar_definition, [new_symbols(naming, [spl150_979])])).
fof(f18527, plain, (~ (e20 = h9(e13)) | ~ spl150_307), inference(backward_demodulation, [], [f5309, f4024])).
fof(f5309, plain, ~ (op2(e24, e23) = h9(e13)), inference(backward_demodulation, [], [f1382, f1534])).
fof(f1382, plain, ~ (op2(e24, e21) = op2(e24, e23)), inference(cnf_transformation, [], [f18])).
fof(f18532, plain, (spl150_785 | ~ spl150_307), inference(avatar_split_clause, [], [f18525, f4022, f7088])).
fof(f18525, plain, ((e20 = h19(e13)) | ~ spl150_307), inference(backward_demodulation, [], [f1594, f4024])).
fof(f18520, plain, (spl150_890 | ~ spl150_313), inference(avatar_split_clause, [], [f18514, f4047, f7643])).
fof(f7643, plain, (spl150_890 <=> (e20 = h14(e13))), introduced(avatar_definition, [new_symbols(naming, [spl150_890])])).
fof(f4047, plain, (spl150_313 <=> (e20 = op2(e24, e22))), introduced(avatar_definition, [new_symbols(naming, [spl150_313])])).
fof(f18514, plain, ((e20 = h14(e13)) | ~ spl150_313), inference(backward_demodulation, [], [f1564, f4049])).
fof(f4049, plain, ((e20 = op2(e24, e22)) | ~ spl150_313), inference(avatar_component_clause, [], [f4047])).
fof(f18510, plain, (spl150_979 | ~ spl150_319), inference(avatar_split_clause, [], [f18505, f4072, f8131])).
fof(f18505, plain, ((e20 = h9(e13)) | ~ spl150_319), inference(backward_demodulation, [], [f1534, f4074])).
fof(f4074, plain, ((e20 = op2(e24, e21)) | ~ spl150_319), inference(avatar_component_clause, [], [f4072])).
fof(f18396, plain, (spl150_624 | ~ spl150_403), inference(avatar_split_clause, [], [f18390, f4422, f6219])).
fof(f6219, plain, (spl150_624 <=> (e20 = h27(e13))), introduced(avatar_definition, [new_symbols(naming, [spl150_624])])).
fof(f4422, plain, (spl150_403 <=> (e20 = op2(e21, e25))), introduced(avatar_definition, [new_symbols(naming, [spl150_403])])).
fof(f18390, plain, ((e20 = h27(e13)) | ~ spl150_403), inference(backward_demodulation, [], [f1642, f4424])).
fof(f4424, plain, ((e20 = op2(e21, e25)) | ~ spl150_403), inference(avatar_component_clause, [], [f4422])).
fof(f18385, plain, (spl150_720 | ~ spl150_409), inference(avatar_split_clause, [], [f18379, f4447, f6745])).
fof(f6745, plain, (spl150_720 <=> (e20 = h22(e13))), introduced(avatar_definition, [new_symbols(naming, [spl150_720])])).
fof(f4447, plain, (spl150_409 <=> (e20 = op2(e21, e24))), introduced(avatar_definition, [new_symbols(naming, [spl150_409])])).
fof(f18379, plain, ((e20 = h22(e13)) | ~ spl150_409), inference(backward_demodulation, [], [f1612, f4449])).
fof(f4449, plain, ((e20 = op2(e21, e24)) | ~ spl150_409), inference(avatar_component_clause, [], [f4447])).
fof(f18335, plain, (spl150_259 | ~ spl150_475 | ~ spl150_486), inference(avatar_split_clause, [], [f18291, f4767, f4722, f3822])).
fof(f18291, plain, ((e20 = op2(e25, e25)) | (~ spl150_475 | ~ spl150_486)), inference(backward_demodulation, [], [f18207, f4724])).
fof(f18322, plain, (~ spl150_890 | ~ spl150_475 | ~ spl150_501), inference(avatar_split_clause, [], [f18278, f4830, f4722, f7643])).
fof(f4830, plain, (spl150_501 <=> (e22 = inv2(e22))), introduced(avatar_definition, [new_symbols(naming, [spl150_501])])).
fof(f18278, plain, (~ (e20 = h14(e13)) | (~ spl150_475 | ~ spl150_501)), inference(backward_demodulation, [], [f18105, f4724])).
fof(f18105, plain, (~ (unit2 = h14(e13)) | ~ spl150_501), inference(backward_demodulation, [], [f5452, f18084])).
fof(f18084, plain, ((op2(e22, e22) = unit2) | ~ spl150_501), inference(backward_demodulation, [], [f916, f4832])).
fof(f4832, plain, ((e22 = inv2(e22)) | ~ spl150_501), inference(avatar_component_clause, [], [f4830])).
fof(f916, plain, (unit2 = op2(e22, inv2(e22))), inference(cnf_transformation, [], [f8])).
fof(f5452, plain, ~ (op2(e22, e22) = h14(e13)), inference(backward_demodulation, [], [f1266, f1564])).
fof(f1266, plain, ~ (op2(e22, e22) = op2(e24, e22)), inference(cnf_transformation, [], [f18])).
fof(f18319, plain, (~ spl150_624 | ~ spl150_475 | ~ spl150_506), inference(avatar_split_clause, [], [f18275, f4851, f4722, f6219])).
fof(f18275, plain, (~ (e20 = h27(e13)) | (~ spl150_475 | ~ spl150_506)), inference(backward_demodulation, [], [f18068, f4724])).
fof(f18068, plain, (~ (unit2 = h27(e13)) | ~ spl150_506), inference(backward_demodulation, [], [f5734, f18037])).
fof(f5734, plain, ~ (op2(e21, e21) = h27(e13)), inference(backward_demodulation, [], [f1339, f1642])).
fof(f1339, plain, ~ (op2(e21, e21) = op2(e21, e25)), inference(cnf_transformation, [], [f18])).
fof(f18318, plain, (~ spl150_720 | ~ spl150_475 | ~ spl150_506), inference(avatar_split_clause, [], [f18274, f4851, f4722, f6745])).
fof(f18274, plain, (~ (e20 = h22(e13)) | (~ spl150_475 | ~ spl150_506)), inference(backward_demodulation, [], [f18066, f4724])).
fof(f18066, plain, (~ (unit2 = h22(e13)) | ~ spl150_506), inference(backward_demodulation, [], [f5634, f18037])).
fof(f5634, plain, ~ (op2(e21, e21) = h22(e13)), inference(backward_demodulation, [], [f1338, f1612])).
fof(f1338, plain, ~ (op2(e21, e21) = op2(e21, e24)), inference(cnf_transformation, [], [f18])).
fof(f18304, plain, (spl150_294 | ~ spl150_475), inference(avatar_split_clause, [], [f18254, f4722, f3967])).
fof(f18254, plain, ((e25 = op2(e25, e20)) | ~ spl150_475), inference(backward_demodulation, [], [f910, f4724])).
fof(f18303, plain, (spl150_444 | ~ spl150_475), inference(avatar_split_clause, [], [f18253, f4722, f4592])).
fof(f18253, plain, ((e25 = op2(e20, e25)) | ~ spl150_475), inference(backward_demodulation, [], [f909, f4724])).
fof(f909, plain, (e25 = op2(unit2, e25)), inference(cnf_transformation, [], [f7])).
fof(f18302, plain, (spl150_329 | ~ spl150_475), inference(avatar_split_clause, [], [f18252, f4722, f4113])).
fof(f18252, plain, ((e24 = op2(e24, e20)) | ~ spl150_475), inference(backward_demodulation, [], [f908, f4724])).
fof(f908, plain, (e24 = op2(e24, unit2)), inference(cnf_transformation, [], [f7])).
fof(f18301, plain, (spl150_449 | ~ spl150_475), inference(avatar_split_clause, [], [f18251, f4722, f4613])).
fof(f18251, plain, ((e24 = op2(e20, e24)) | ~ spl150_475), inference(backward_demodulation, [], [f907, f4724])).
fof(f18299, plain, (spl150_454 | ~ spl150_475), inference(avatar_split_clause, [], [f18249, f4722, f4634])).
fof(f18249, plain, ((e23 = op2(e20, e23)) | ~ spl150_475), inference(backward_demodulation, [], [f905, f4724])).
fof(f905, plain, (e23 = op2(unit2, e23)), inference(cnf_transformation, [], [f7])).
fof(f18297, plain, (spl150_459 | ~ spl150_475), inference(avatar_split_clause, [], [f18247, f4722, f4655])).
fof(f18247, plain, ((e22 = op2(e20, e22)) | ~ spl150_475), inference(backward_demodulation, [], [f903, f4724])).
fof(f18036, plain, (~ spl150_481 | ~ spl150_511), inference(avatar_split_clause, [], [f17994, f4872, f4747])).
fof(f4747, plain, (spl150_481 <=> (e20 = inv2(e25))), introduced(avatar_definition, [new_symbols(naming, [spl150_481])])).
fof(f17994, plain, (~ (e20 = inv2(e25)) | ~ spl150_511), inference(backward_demodulation, [], [f1035, f4874])).
fof(f1035, plain, ~ (inv2(e20) = inv2(e25)), inference(cnf_transformation, [], [f16])).
fof(f18034, plain, (~ spl150_493 | ~ spl150_511), inference(avatar_split_clause, [], [f17992, f4872, f4797])).
fof(f4797, plain, (spl150_493 <=> (e20 = inv2(e23))), introduced(avatar_definition, [new_symbols(naming, [spl150_493])])).
fof(f17992, plain, (~ (e20 = inv2(e23)) | ~ spl150_511), inference(backward_demodulation, [], [f1033, f4874])).
fof(f1033, plain, ~ (inv2(e20) = inv2(e23)), inference(cnf_transformation, [], [f16])).
fof(f18033, plain, (~ spl150_499 | ~ spl150_511), inference(avatar_split_clause, [], [f17991, f4872, f4822])).
fof(f4822, plain, (spl150_499 <=> (e20 = inv2(e22))), introduced(avatar_definition, [new_symbols(naming, [spl150_499])])).
fof(f17991, plain, (~ (e20 = inv2(e22)) | ~ spl150_511), inference(backward_demodulation, [], [f1032, f4874])).
fof(f1032, plain, ~ (inv2(e20) = inv2(e22)), inference(cnf_transformation, [], [f16])).
fof(f15272, plain, (~ spl150_1938 | ~ spl150_1939 | spl150_959 | spl150_957 | spl150_955 | spl150_953 | ~ spl150_1940 | ~ spl150_1941 | ~ spl150_1942 | ~ spl150_1943 | ~ spl150_1944 | ~ spl150_1945 | ~ spl150_1946 | ~ spl150_1947 | ~ spl150_1948 | ~ spl150_1949 | ~ spl150_1950 | ~ spl150_1951 | ~ spl150_1952 | ~ spl150_1953 | ~ spl150_1954 | ~ spl150_1955 | ~ spl150_1956 | ~ spl150_1957 | ~ spl150_1958 | ~ spl150_1959 | ~ spl150_1960 | ~ spl150_1961 | ~ spl150_1962 | ~ spl150_1963 | ~ spl150_1964 | ~ spl150_1965 | ~ spl150_1966 | ~ spl150_1967 | ~ spl150_1968 | ~ spl150_1969 | ~ spl150_1970 | ~ spl150_1971 | ~ spl150_1972 | ~ spl150_1973), inference(avatar_split_clause, [], [f15127, f15269, f15265, f15261, f15257, f15253, f15249, f15245, f15241, f15237, f15233, f15229, f15225, f15221, f15217, f15213, f15209, f15205, f15201, f15197, f15193, f15189, f15185, f15181, f15177, f15173, f15169, f15165, f15161, f15157, f15153, f15149, f15145, f15141, f15137, f7981, f7995, f8009, f8024, f15133, f15129])).
fof(f8024, plain, (spl150_959 <=> sP45), introduced(avatar_definition, [new_symbols(naming, [spl150_959])])).
fof(f8009, plain, (spl150_957 <=> sP46), introduced(avatar_definition, [new_symbols(naming, [spl150_957])])).
fof(f7995, plain, (spl150_955 <=> sP47), introduced(avatar_definition, [new_symbols(naming, [spl150_955])])).
fof(f7981, plain, (spl150_953 <=> sP48), introduced(avatar_definition, [new_symbols(naming, [spl150_953])])).
fof(f15127, plain, (~ (op2(e20, e20) = h10(op1(e10, e10))) | ~ (h6(e13) = h10(op1(e10, e11))) | ~ (h10(op1(e10, e12)) = op2(e20, h10(e12))) | ~ (h10(op1(e10, e13)) = op2(e20, h10(e13))) | ~ (h21(e13) = h10(op1(e10, e14))) | ~ (h26(e13) = h10(op1(e10, e15))) | ~ (h1(e13) = h10(op1(e11, e10))) | ~ (op2(e21, e21) = h10(op1(e11, e11))) | ~ (h10(op1(e11, e12)) = op2(e21, h10(e12))) | ~ (h10(op1(e11, e13)) = op2(e21, h10(e13))) | ~ (h22(e13) = h10(op1(e11, e14))) | ~ (h27(e13) = h10(op1(e11, e15))) | ~ (h10(op1(e12, e10)) = op2(e25, h4(e13))) | ~ (h10(op1(e12, e11)) = op2(e25, h9(e13))) | ~ (e20 = h10(op1(e12, e12))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), e24)) | ~ (h10(op1(e12, e15)) = op2(e25, h30(e13))) | ~ (h10(op1(e13, e10)) = op2(e25, h1(e13))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), e21)) | ~ (e24 = h10(op1(e13, e13))) | ~ (h10(op1(e13, e14)) = op2(e25, h22(e13))) | ~ (h10(op1(e13, e15)) = op2(e25, h27(e13))) | ~ (h4(e13) = h10(op1(e14, e10))) | ~ (h9(e13) = h10(op1(e14, e11))) | ~ (h10(op1(e14, e12)) = op2(e24, h10(e12))) | ~ (h10(op1(e14, e13)) = op2(e24, h10(e13))) | ~ (op2(e24, e24) = h10(op1(e14, e14))) | ~ (h30(e13) = h10(op1(e14, e15))) | ~ (h5(e13) = h10(op1(e15, e10))) | ~ (h10(e13) = h10(op1(e15, e11))) | ~ (h10(op1(e15, e12)) = op2(e25, h10(e12))) | ~ (h10(op1(e15, e13)) = op2(e25, h10(e13))) | ~ (h10(e12) = h10(op1(e15, e14))) | ~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13)))), inference(forward_demodulation, [], [f15126, f5365])).
fof(f15126, plain, (~ (h6(e13) = h10(op1(e10, e11))) | ~ (h10(op1(e10, e12)) = op2(e20, h10(e12))) | ~ (h10(op1(e10, e13)) = op2(e20, h10(e13))) | ~ (h21(e13) = h10(op1(e10, e14))) | ~ (h26(e13) = h10(op1(e10, e15))) | ~ (h1(e13) = h10(op1(e11, e10))) | ~ (op2(e21, e21) = h10(op1(e11, e11))) | ~ (h10(op1(e11, e12)) = op2(e21, h10(e12))) | ~ (h10(op1(e11, e13)) = op2(e21, h10(e13))) | ~ (h22(e13) = h10(op1(e11, e14))) | ~ (h27(e13) = h10(op1(e11, e15))) | ~ (h10(op1(e12, e10)) = op2(e25, h4(e13))) | ~ (h10(op1(e12, e11)) = op2(e25, h9(e13))) | ~ (e20 = h10(op1(e12, e12))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), e24)) | ~ (h10(op1(e12, e15)) = op2(e25, h30(e13))) | ~ (h10(op1(e13, e10)) = op2(e25, h1(e13))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), e21)) | ~ (e24 = h10(op1(e13, e13))) | ~ (h10(op1(e13, e14)) = op2(e25, h22(e13))) | ~ (h10(op1(e13, e15)) = op2(e25, h27(e13))) | ~ (h4(e13) = h10(op1(e14, e10))) | ~ (h9(e13) = h10(op1(e14, e11))) | ~ (h10(op1(e14, e12)) = op2(e24, h10(e12))) | ~ (h10(op1(e14, e13)) = op2(e24, h10(e13))) | ~ (op2(e24, e24) = h10(op1(e14, e14))) | ~ (h30(e13) = h10(op1(e14, e15))) | ~ (h5(e13) = h10(op1(e15, e10))) | ~ (h10(e13) = h10(op1(e15, e11))) | ~ (h10(op1(e15, e12)) = op2(e25, h10(e12))) | ~ (h10(op1(e15, e13)) = op2(e25, h10(e13))) | ~ (h10(e12) = h10(op1(e15, e14))) | ~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(forward_demodulation, [], [f15125, f1516])).
fof(f15125, plain, (~ (op2(e20, e21) = h10(op1(e10, e11))) | ~ (h10(op1(e10, e12)) = op2(e20, h10(e12))) | ~ (h10(op1(e10, e13)) = op2(e20, h10(e13))) | ~ (h21(e13) = h10(op1(e10, e14))) | ~ (h26(e13) = h10(op1(e10, e15))) | ~ (h1(e13) = h10(op1(e11, e10))) | ~ (op2(e21, e21) = h10(op1(e11, e11))) | ~ (h10(op1(e11, e12)) = op2(e21, h10(e12))) | ~ (h10(op1(e11, e13)) = op2(e21, h10(e13))) | ~ (h22(e13) = h10(op1(e11, e14))) | ~ (h27(e13) = h10(op1(e11, e15))) | ~ (h10(op1(e12, e10)) = op2(e25, h4(e13))) | ~ (h10(op1(e12, e11)) = op2(e25, h9(e13))) | ~ (e20 = h10(op1(e12, e12))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), e24)) | ~ (h10(op1(e12, e15)) = op2(e25, h30(e13))) | ~ (h10(op1(e13, e10)) = op2(e25, h1(e13))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), e21)) | ~ (e24 = h10(op1(e13, e13))) | ~ (h10(op1(e13, e14)) = op2(e25, h22(e13))) | ~ (h10(op1(e13, e15)) = op2(e25, h27(e13))) | ~ (h4(e13) = h10(op1(e14, e10))) | ~ (h9(e13) = h10(op1(e14, e11))) | ~ (h10(op1(e14, e12)) = op2(e24, h10(e12))) | ~ (h10(op1(e14, e13)) = op2(e24, h10(e13))) | ~ (op2(e24, e24) = h10(op1(e14, e14))) | ~ (h30(e13) = h10(op1(e14, e15))) | ~ (h5(e13) = h10(op1(e15, e10))) | ~ (h10(e13) = h10(op1(e15, e11))) | ~ (h10(op1(e15, e12)) = op2(e25, h10(e12))) | ~ (h10(op1(e15, e13)) = op2(e25, h10(e13))) | ~ (h10(e12) = h10(op1(e15, e14))) | ~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(forward_demodulation, [], [f15124, f5365])).
fof(f15124, plain, (~ (h10(op1(e10, e11)) = op2(h10(e10), e21)) | ~ (h10(op1(e10, e12)) = op2(e20, h10(e12))) | ~ (h10(op1(e10, e13)) = op2(e20, h10(e13))) | ~ (h21(e13) = h10(op1(e10, e14))) | ~ (h26(e13) = h10(op1(e10, e15))) | ~ (h1(e13) = h10(op1(e11, e10))) | ~ (op2(e21, e21) = h10(op1(e11, e11))) | ~ (h10(op1(e11, e12)) = op2(e21, h10(e12))) | ~ (h10(op1(e11, e13)) = op2(e21, h10(e13))) | ~ (h22(e13) = h10(op1(e11, e14))) | ~ (h27(e13) = h10(op1(e11, e15))) | ~ (h10(op1(e12, e10)) = op2(e25, h4(e13))) | ~ (h10(op1(e12, e11)) = op2(e25, h9(e13))) | ~ (e20 = h10(op1(e12, e12))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), e24)) | ~ (h10(op1(e12, e15)) = op2(e25, h30(e13))) | ~ (h10(op1(e13, e10)) = op2(e25, h1(e13))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), e21)) | ~ (e24 = h10(op1(e13, e13))) | ~ (h10(op1(e13, e14)) = op2(e25, h22(e13))) | ~ (h10(op1(e13, e15)) = op2(e25, h27(e13))) | ~ (h4(e13) = h10(op1(e14, e10))) | ~ (h9(e13) = h10(op1(e14, e11))) | ~ (h10(op1(e14, e12)) = op2(e24, h10(e12))) | ~ (h10(op1(e14, e13)) = op2(e24, h10(e13))) | ~ (op2(e24, e24) = h10(op1(e14, e14))) | ~ (h30(e13) = h10(op1(e14, e15))) | ~ (h5(e13) = h10(op1(e15, e10))) | ~ (h10(e13) = h10(op1(e15, e11))) | ~ (h10(op1(e15, e12)) = op2(e25, h10(e12))) | ~ (h10(op1(e15, e13)) = op2(e25, h10(e13))) | ~ (h10(e12) = h10(op1(e15, e14))) | ~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(forward_demodulation, [], [f15123, f1536])).
fof(f15123, plain, (~ (h10(op1(e10, e12)) = op2(e20, h10(e12))) | ~ (h10(op1(e10, e13)) = op2(e20, h10(e13))) | ~ (h21(e13) = h10(op1(e10, e14))) | ~ (h26(e13) = h10(op1(e10, e15))) | ~ (h1(e13) = h10(op1(e11, e10))) | ~ (op2(e21, e21) = h10(op1(e11, e11))) | ~ (h10(op1(e11, e12)) = op2(e21, h10(e12))) | ~ (h10(op1(e11, e13)) = op2(e21, h10(e13))) | ~ (h22(e13) = h10(op1(e11, e14))) | ~ (h27(e13) = h10(op1(e11, e15))) | ~ (h10(op1(e12, e10)) = op2(e25, h4(e13))) | ~ (h10(op1(e12, e11)) = op2(e25, h9(e13))) | ~ (e20 = h10(op1(e12, e12))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), e24)) | ~ (h10(op1(e12, e15)) = op2(e25, h30(e13))) | ~ (h10(op1(e13, e10)) = op2(e25, h1(e13))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), e21)) | ~ (e24 = h10(op1(e13, e13))) | ~ (h10(op1(e13, e14)) = op2(e25, h22(e13))) | ~ (h10(op1(e13, e15)) = op2(e25, h27(e13))) | ~ (h4(e13) = h10(op1(e14, e10))) | ~ (h9(e13) = h10(op1(e14, e11))) | ~ (h10(op1(e14, e12)) = op2(e24, h10(e12))) | ~ (h10(op1(e14, e13)) = op2(e24, h10(e13))) | ~ (op2(e24, e24) = h10(op1(e14, e14))) | ~ (h30(e13) = h10(op1(e14, e15))) | ~ (h5(e13) = h10(op1(e15, e10))) | ~ (h10(e13) = h10(op1(e15, e11))) | ~ (h10(op1(e15, e12)) = op2(e25, h10(e12))) | ~ (h10(op1(e15, e13)) = op2(e25, h10(e13))) | ~ (h10(e12) = h10(op1(e15, e14))) | ~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(forward_demodulation, [], [f15122, f5365])).
fof(f15122, plain, (~ (h10(op1(e10, e13)) = op2(e20, h10(e13))) | ~ (h21(e13) = h10(op1(e10, e14))) | ~ (h26(e13) = h10(op1(e10, e15))) | ~ (h1(e13) = h10(op1(e11, e10))) | ~ (op2(e21, e21) = h10(op1(e11, e11))) | ~ (h10(op1(e11, e12)) = op2(e21, h10(e12))) | ~ (h10(op1(e11, e13)) = op2(e21, h10(e13))) | ~ (h22(e13) = h10(op1(e11, e14))) | ~ (h27(e13) = h10(op1(e11, e15))) | ~ (h10(op1(e12, e10)) = op2(e25, h4(e13))) | ~ (h10(op1(e12, e11)) = op2(e25, h9(e13))) | ~ (e20 = h10(op1(e12, e12))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), e24)) | ~ (h10(op1(e12, e15)) = op2(e25, h30(e13))) | ~ (h10(op1(e13, e10)) = op2(e25, h1(e13))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), e21)) | ~ (e24 = h10(op1(e13, e13))) | ~ (h10(op1(e13, e14)) = op2(e25, h22(e13))) | ~ (h10(op1(e13, e15)) = op2(e25, h27(e13))) | ~ (h4(e13) = h10(op1(e14, e10))) | ~ (h9(e13) = h10(op1(e14, e11))) | ~ (h10(op1(e14, e12)) = op2(e24, h10(e12))) | ~ (h10(op1(e14, e13)) = op2(e24, h10(e13))) | ~ (op2(e24, e24) = h10(op1(e14, e14))) | ~ (h30(e13) = h10(op1(e14, e15))) | ~ (h5(e13) = h10(op1(e15, e10))) | ~ (h10(e13) = h10(op1(e15, e11))) | ~ (h10(op1(e15, e12)) = op2(e25, h10(e12))) | ~ (h10(op1(e15, e13)) = op2(e25, h10(e13))) | ~ (h10(e12) = h10(op1(e15, e14))) | ~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(forward_demodulation, [], [f15121, f5365])).
fof(f15121, plain, (~ (h21(e13) = h10(op1(e10, e14))) | ~ (h26(e13) = h10(op1(e10, e15))) | ~ (h1(e13) = h10(op1(e11, e10))) | ~ (op2(e21, e21) = h10(op1(e11, e11))) | ~ (h10(op1(e11, e12)) = op2(e21, h10(e12))) | ~ (h10(op1(e11, e13)) = op2(e21, h10(e13))) | ~ (h22(e13) = h10(op1(e11, e14))) | ~ (h27(e13) = h10(op1(e11, e15))) | ~ (h10(op1(e12, e10)) = op2(e25, h4(e13))) | ~ (h10(op1(e12, e11)) = op2(e25, h9(e13))) | ~ (e20 = h10(op1(e12, e12))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), e24)) | ~ (h10(op1(e12, e15)) = op2(e25, h30(e13))) | ~ (h10(op1(e13, e10)) = op2(e25, h1(e13))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), e21)) | ~ (e24 = h10(op1(e13, e13))) | ~ (h10(op1(e13, e14)) = op2(e25, h22(e13))) | ~ (h10(op1(e13, e15)) = op2(e25, h27(e13))) | ~ (h4(e13) = h10(op1(e14, e10))) | ~ (h9(e13) = h10(op1(e14, e11))) | ~ (h10(op1(e14, e12)) = op2(e24, h10(e12))) | ~ (h10(op1(e14, e13)) = op2(e24, h10(e13))) | ~ (op2(e24, e24) = h10(op1(e14, e14))) | ~ (h30(e13) = h10(op1(e14, e15))) | ~ (h5(e13) = h10(op1(e15, e10))) | ~ (h10(e13) = h10(op1(e15, e11))) | ~ (h10(op1(e15, e12)) = op2(e25, h10(e12))) | ~ (h10(op1(e15, e13)) = op2(e25, h10(e13))) | ~ (h10(e12) = h10(op1(e15, e14))) | ~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(forward_demodulation, [], [f15120, f1606])).
fof(f15120, plain, (~ (op2(e20, e24) = h10(op1(e10, e14))) | ~ (h26(e13) = h10(op1(e10, e15))) | ~ (h1(e13) = h10(op1(e11, e10))) | ~ (op2(e21, e21) = h10(op1(e11, e11))) | ~ (h10(op1(e11, e12)) = op2(e21, h10(e12))) | ~ (h10(op1(e11, e13)) = op2(e21, h10(e13))) | ~ (h22(e13) = h10(op1(e11, e14))) | ~ (h27(e13) = h10(op1(e11, e15))) | ~ (h10(op1(e12, e10)) = op2(e25, h4(e13))) | ~ (h10(op1(e12, e11)) = op2(e25, h9(e13))) | ~ (e20 = h10(op1(e12, e12))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), e24)) | ~ (h10(op1(e12, e15)) = op2(e25, h30(e13))) | ~ (h10(op1(e13, e10)) = op2(e25, h1(e13))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), e21)) | ~ (e24 = h10(op1(e13, e13))) | ~ (h10(op1(e13, e14)) = op2(e25, h22(e13))) | ~ (h10(op1(e13, e15)) = op2(e25, h27(e13))) | ~ (h4(e13) = h10(op1(e14, e10))) | ~ (h9(e13) = h10(op1(e14, e11))) | ~ (h10(op1(e14, e12)) = op2(e24, h10(e12))) | ~ (h10(op1(e14, e13)) = op2(e24, h10(e13))) | ~ (op2(e24, e24) = h10(op1(e14, e14))) | ~ (h30(e13) = h10(op1(e14, e15))) | ~ (h5(e13) = h10(op1(e15, e10))) | ~ (h10(e13) = h10(op1(e15, e11))) | ~ (h10(op1(e15, e12)) = op2(e25, h10(e12))) | ~ (h10(op1(e15, e13)) = op2(e25, h10(e13))) | ~ (h10(e12) = h10(op1(e15, e14))) | ~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(forward_demodulation, [], [f15119, f5365])).
fof(f15119, plain, (~ (h10(op1(e10, e14)) = op2(h10(e10), e24)) | ~ (h26(e13) = h10(op1(e10, e15))) | ~ (h1(e13) = h10(op1(e11, e10))) | ~ (op2(e21, e21) = h10(op1(e11, e11))) | ~ (h10(op1(e11, e12)) = op2(e21, h10(e12))) | ~ (h10(op1(e11, e13)) = op2(e21, h10(e13))) | ~ (h22(e13) = h10(op1(e11, e14))) | ~ (h27(e13) = h10(op1(e11, e15))) | ~ (h10(op1(e12, e10)) = op2(e25, h4(e13))) | ~ (h10(op1(e12, e11)) = op2(e25, h9(e13))) | ~ (e20 = h10(op1(e12, e12))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), e24)) | ~ (h10(op1(e12, e15)) = op2(e25, h30(e13))) | ~ (h10(op1(e13, e10)) = op2(e25, h1(e13))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), e21)) | ~ (e24 = h10(op1(e13, e13))) | ~ (h10(op1(e13, e14)) = op2(e25, h22(e13))) | ~ (h10(op1(e13, e15)) = op2(e25, h27(e13))) | ~ (h4(e13) = h10(op1(e14, e10))) | ~ (h9(e13) = h10(op1(e14, e11))) | ~ (h10(op1(e14, e12)) = op2(e24, h10(e12))) | ~ (h10(op1(e14, e13)) = op2(e24, h10(e13))) | ~ (op2(e24, e24) = h10(op1(e14, e14))) | ~ (h30(e13) = h10(op1(e14, e15))) | ~ (h5(e13) = h10(op1(e15, e10))) | ~ (h10(e13) = h10(op1(e15, e11))) | ~ (h10(op1(e15, e12)) = op2(e25, h10(e12))) | ~ (h10(op1(e15, e13)) = op2(e25, h10(e13))) | ~ (h10(e12) = h10(op1(e15, e14))) | ~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(forward_demodulation, [], [f15118, f5319])).
fof(f15118, plain, (~ (h26(e13) = h10(op1(e10, e15))) | ~ (h1(e13) = h10(op1(e11, e10))) | ~ (op2(e21, e21) = h10(op1(e11, e11))) | ~ (h10(op1(e11, e12)) = op2(e21, h10(e12))) | ~ (h10(op1(e11, e13)) = op2(e21, h10(e13))) | ~ (h22(e13) = h10(op1(e11, e14))) | ~ (h27(e13) = h10(op1(e11, e15))) | ~ (h10(op1(e12, e10)) = op2(e25, h4(e13))) | ~ (h10(op1(e12, e11)) = op2(e25, h9(e13))) | ~ (e20 = h10(op1(e12, e12))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), e24)) | ~ (h10(op1(e12, e15)) = op2(e25, h30(e13))) | ~ (h10(op1(e13, e10)) = op2(e25, h1(e13))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), e21)) | ~ (e24 = h10(op1(e13, e13))) | ~ (h10(op1(e13, e14)) = op2(e25, h22(e13))) | ~ (h10(op1(e13, e15)) = op2(e25, h27(e13))) | ~ (h4(e13) = h10(op1(e14, e10))) | ~ (h9(e13) = h10(op1(e14, e11))) | ~ (h10(op1(e14, e12)) = op2(e24, h10(e12))) | ~ (h10(op1(e14, e13)) = op2(e24, h10(e13))) | ~ (op2(e24, e24) = h10(op1(e14, e14))) | ~ (h30(e13) = h10(op1(e14, e15))) | ~ (h5(e13) = h10(op1(e15, e10))) | ~ (h10(e13) = h10(op1(e15, e11))) | ~ (h10(op1(e15, e12)) = op2(e25, h10(e12))) | ~ (h10(op1(e15, e13)) = op2(e25, h10(e13))) | ~ (h10(e12) = h10(op1(e15, e14))) | ~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(forward_demodulation, [], [f15117, f1636])).
fof(f15117, plain, (~ (op2(e20, e25) = h10(op1(e10, e15))) | ~ (h1(e13) = h10(op1(e11, e10))) | ~ (op2(e21, e21) = h10(op1(e11, e11))) | ~ (h10(op1(e11, e12)) = op2(e21, h10(e12))) | ~ (h10(op1(e11, e13)) = op2(e21, h10(e13))) | ~ (h22(e13) = h10(op1(e11, e14))) | ~ (h27(e13) = h10(op1(e11, e15))) | ~ (h10(op1(e12, e10)) = op2(e25, h4(e13))) | ~ (h10(op1(e12, e11)) = op2(e25, h9(e13))) | ~ (e20 = h10(op1(e12, e12))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), e24)) | ~ (h10(op1(e12, e15)) = op2(e25, h30(e13))) | ~ (h10(op1(e13, e10)) = op2(e25, h1(e13))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), e21)) | ~ (e24 = h10(op1(e13, e13))) | ~ (h10(op1(e13, e14)) = op2(e25, h22(e13))) | ~ (h10(op1(e13, e15)) = op2(e25, h27(e13))) | ~ (h4(e13) = h10(op1(e14, e10))) | ~ (h9(e13) = h10(op1(e14, e11))) | ~ (h10(op1(e14, e12)) = op2(e24, h10(e12))) | ~ (h10(op1(e14, e13)) = op2(e24, h10(e13))) | ~ (op2(e24, e24) = h10(op1(e14, e14))) | ~ (h30(e13) = h10(op1(e14, e15))) | ~ (h5(e13) = h10(op1(e15, e10))) | ~ (h10(e13) = h10(op1(e15, e11))) | ~ (h10(op1(e15, e12)) = op2(e25, h10(e12))) | ~ (h10(op1(e15, e13)) = op2(e25, h10(e13))) | ~ (h10(e12) = h10(op1(e15, e14))) | ~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(forward_demodulation, [], [f15116, f5365])).
fof(f15116, plain, (~ (h10(op1(e10, e15)) = op2(h10(e10), e25)) | ~ (h1(e13) = h10(op1(e11, e10))) | ~ (op2(e21, e21) = h10(op1(e11, e11))) | ~ (h10(op1(e11, e12)) = op2(e21, h10(e12))) | ~ (h10(op1(e11, e13)) = op2(e21, h10(e13))) | ~ (h22(e13) = h10(op1(e11, e14))) | ~ (h27(e13) = h10(op1(e11, e15))) | ~ (h10(op1(e12, e10)) = op2(e25, h4(e13))) | ~ (h10(op1(e12, e11)) = op2(e25, h9(e13))) | ~ (e20 = h10(op1(e12, e12))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), e24)) | ~ (h10(op1(e12, e15)) = op2(e25, h30(e13))) | ~ (h10(op1(e13, e10)) = op2(e25, h1(e13))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), e21)) | ~ (e24 = h10(op1(e13, e13))) | ~ (h10(op1(e13, e14)) = op2(e25, h22(e13))) | ~ (h10(op1(e13, e15)) = op2(e25, h27(e13))) | ~ (h4(e13) = h10(op1(e14, e10))) | ~ (h9(e13) = h10(op1(e14, e11))) | ~ (h10(op1(e14, e12)) = op2(e24, h10(e12))) | ~ (h10(op1(e14, e13)) = op2(e24, h10(e13))) | ~ (op2(e24, e24) = h10(op1(e14, e14))) | ~ (h30(e13) = h10(op1(e14, e15))) | ~ (h5(e13) = h10(op1(e15, e10))) | ~ (h10(e13) = h10(op1(e15, e11))) | ~ (h10(op1(e15, e12)) = op2(e25, h10(e12))) | ~ (h10(op1(e15, e13)) = op2(e25, h10(e13))) | ~ (h10(e12) = h10(op1(e15, e14))) | ~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(forward_demodulation, [], [f15115, f1537])).
fof(f15115, plain, (~ (h1(e13) = h10(op1(e11, e10))) | ~ (op2(e21, e21) = h10(op1(e11, e11))) | ~ (h10(op1(e11, e12)) = op2(e21, h10(e12))) | ~ (h10(op1(e11, e13)) = op2(e21, h10(e13))) | ~ (h22(e13) = h10(op1(e11, e14))) | ~ (h27(e13) = h10(op1(e11, e15))) | ~ (h10(op1(e12, e10)) = op2(e25, h4(e13))) | ~ (h10(op1(e12, e11)) = op2(e25, h9(e13))) | ~ (e20 = h10(op1(e12, e12))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), e24)) | ~ (h10(op1(e12, e15)) = op2(e25, h30(e13))) | ~ (h10(op1(e13, e10)) = op2(e25, h1(e13))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), e21)) | ~ (e24 = h10(op1(e13, e13))) | ~ (h10(op1(e13, e14)) = op2(e25, h22(e13))) | ~ (h10(op1(e13, e15)) = op2(e25, h27(e13))) | ~ (h4(e13) = h10(op1(e14, e10))) | ~ (h9(e13) = h10(op1(e14, e11))) | ~ (h10(op1(e14, e12)) = op2(e24, h10(e12))) | ~ (h10(op1(e14, e13)) = op2(e24, h10(e13))) | ~ (op2(e24, e24) = h10(op1(e14, e14))) | ~ (h30(e13) = h10(op1(e14, e15))) | ~ (h5(e13) = h10(op1(e15, e10))) | ~ (h10(e13) = h10(op1(e15, e11))) | ~ (h10(op1(e15, e12)) = op2(e25, h10(e12))) | ~ (h10(op1(e15, e13)) = op2(e25, h10(e13))) | ~ (h10(e12) = h10(op1(e15, e14))) | ~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e10, e15)) = op2(h10(e10), h10(e15))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(forward_demodulation, [], [f15114, f1486])).
fof(f15114, plain, (~ (op2(e21, e20) = h10(op1(e11, e10))) | ~ (op2(e21, e21) = h10(op1(e11, e11))) | ~ (h10(op1(e11, e12)) = op2(e21, h10(e12))) | ~ (h10(op1(e11, e13)) = op2(e21, h10(e13))) | ~ (h22(e13) = h10(op1(e11, e14))) | ~ (h27(e13) = h10(op1(e11, e15))) | ~ (h10(op1(e12, e10)) = op2(e25, h4(e13))) | ~ (h10(op1(e12, e11)) = op2(e25, h9(e13))) | ~ (e20 = h10(op1(e12, e12))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), e24)) | ~ (h10(op1(e12, e15)) = op2(e25, h30(e13))) | ~ (h10(op1(e13, e10)) = op2(e25, h1(e13))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), e21)) | ~ (e24 = h10(op1(e13, e13))) | ~ (h10(op1(e13, e14)) = op2(e25, h22(e13))) | ~ (h10(op1(e13, e15)) = op2(e25, h27(e13))) | ~ (h4(e13) = h10(op1(e14, e10))) | ~ (h9(e13) = h10(op1(e14, e11))) | ~ (h10(op1(e14, e12)) = op2(e24, h10(e12))) | ~ (h10(op1(e14, e13)) = op2(e24, h10(e13))) | ~ (op2(e24, e24) = h10(op1(e14, e14))) | ~ (h30(e13) = h10(op1(e14, e15))) | ~ (h5(e13) = h10(op1(e15, e10))) | ~ (h10(e13) = h10(op1(e15, e11))) | ~ (h10(op1(e15, e12)) = op2(e25, h10(e12))) | ~ (h10(op1(e15, e13)) = op2(e25, h10(e13))) | ~ (h10(e12) = h10(op1(e15, e14))) | ~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e10, e15)) = op2(h10(e10), h10(e15))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(forward_demodulation, [], [f15113, f1536])).
fof(f15113, plain, (~ (h10(op1(e11, e10)) = op2(h10(e11), e20)) | ~ (op2(e21, e21) = h10(op1(e11, e11))) | ~ (h10(op1(e11, e12)) = op2(e21, h10(e12))) | ~ (h10(op1(e11, e13)) = op2(e21, h10(e13))) | ~ (h22(e13) = h10(op1(e11, e14))) | ~ (h27(e13) = h10(op1(e11, e15))) | ~ (h10(op1(e12, e10)) = op2(e25, h4(e13))) | ~ (h10(op1(e12, e11)) = op2(e25, h9(e13))) | ~ (e20 = h10(op1(e12, e12))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), e24)) | ~ (h10(op1(e12, e15)) = op2(e25, h30(e13))) | ~ (h10(op1(e13, e10)) = op2(e25, h1(e13))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), e21)) | ~ (e24 = h10(op1(e13, e13))) | ~ (h10(op1(e13, e14)) = op2(e25, h22(e13))) | ~ (h10(op1(e13, e15)) = op2(e25, h27(e13))) | ~ (h4(e13) = h10(op1(e14, e10))) | ~ (h9(e13) = h10(op1(e14, e11))) | ~ (h10(op1(e14, e12)) = op2(e24, h10(e12))) | ~ (h10(op1(e14, e13)) = op2(e24, h10(e13))) | ~ (op2(e24, e24) = h10(op1(e14, e14))) | ~ (h30(e13) = h10(op1(e14, e15))) | ~ (h5(e13) = h10(op1(e15, e10))) | ~ (h10(e13) = h10(op1(e15, e11))) | ~ (h10(op1(e15, e12)) = op2(e25, h10(e12))) | ~ (h10(op1(e15, e13)) = op2(e25, h10(e13))) | ~ (h10(e12) = h10(op1(e15, e14))) | ~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e10, e15)) = op2(h10(e10), h10(e15))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(forward_demodulation, [], [f15112, f5365])).
fof(f15112, plain, (~ (op2(e21, e21) = h10(op1(e11, e11))) | ~ (h10(op1(e11, e12)) = op2(e21, h10(e12))) | ~ (h10(op1(e11, e13)) = op2(e21, h10(e13))) | ~ (h22(e13) = h10(op1(e11, e14))) | ~ (h27(e13) = h10(op1(e11, e15))) | ~ (h10(op1(e12, e10)) = op2(e25, h4(e13))) | ~ (h10(op1(e12, e11)) = op2(e25, h9(e13))) | ~ (e20 = h10(op1(e12, e12))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), e24)) | ~ (h10(op1(e12, e15)) = op2(e25, h30(e13))) | ~ (h10(op1(e13, e10)) = op2(e25, h1(e13))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), e21)) | ~ (e24 = h10(op1(e13, e13))) | ~ (h10(op1(e13, e14)) = op2(e25, h22(e13))) | ~ (h10(op1(e13, e15)) = op2(e25, h27(e13))) | ~ (h4(e13) = h10(op1(e14, e10))) | ~ (h9(e13) = h10(op1(e14, e11))) | ~ (h10(op1(e14, e12)) = op2(e24, h10(e12))) | ~ (h10(op1(e14, e13)) = op2(e24, h10(e13))) | ~ (op2(e24, e24) = h10(op1(e14, e14))) | ~ (h30(e13) = h10(op1(e14, e15))) | ~ (h5(e13) = h10(op1(e15, e10))) | ~ (h10(e13) = h10(op1(e15, e11))) | ~ (h10(op1(e15, e12)) = op2(e25, h10(e12))) | ~ (h10(op1(e15, e13)) = op2(e25, h10(e13))) | ~ (h10(e12) = h10(op1(e15, e14))) | ~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e15)) = op2(h10(e10), h10(e15))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(forward_demodulation, [], [f15111, f1536])).
fof(f15111, plain, (~ (h10(op1(e11, e12)) = op2(e21, h10(e12))) | ~ (h10(op1(e11, e13)) = op2(e21, h10(e13))) | ~ (h22(e13) = h10(op1(e11, e14))) | ~ (h27(e13) = h10(op1(e11, e15))) | ~ (h10(op1(e12, e10)) = op2(e25, h4(e13))) | ~ (h10(op1(e12, e11)) = op2(e25, h9(e13))) | ~ (e20 = h10(op1(e12, e12))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), e24)) | ~ (h10(op1(e12, e15)) = op2(e25, h30(e13))) | ~ (h10(op1(e13, e10)) = op2(e25, h1(e13))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), e21)) | ~ (e24 = h10(op1(e13, e13))) | ~ (h10(op1(e13, e14)) = op2(e25, h22(e13))) | ~ (h10(op1(e13, e15)) = op2(e25, h27(e13))) | ~ (h4(e13) = h10(op1(e14, e10))) | ~ (h9(e13) = h10(op1(e14, e11))) | ~ (h10(op1(e14, e12)) = op2(e24, h10(e12))) | ~ (h10(op1(e14, e13)) = op2(e24, h10(e13))) | ~ (op2(e24, e24) = h10(op1(e14, e14))) | ~ (h30(e13) = h10(op1(e14, e15))) | ~ (h5(e13) = h10(op1(e15, e10))) | ~ (h10(e13) = h10(op1(e15, e11))) | ~ (h10(op1(e15, e12)) = op2(e25, h10(e12))) | ~ (h10(op1(e15, e13)) = op2(e25, h10(e13))) | ~ (h10(e12) = h10(op1(e15, e14))) | ~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e15)) = op2(h10(e10), h10(e15))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(forward_demodulation, [], [f15110, f1536])).
fof(f15110, plain, (~ (h10(op1(e11, e13)) = op2(e21, h10(e13))) | ~ (h22(e13) = h10(op1(e11, e14))) | ~ (h27(e13) = h10(op1(e11, e15))) | ~ (h10(op1(e12, e10)) = op2(e25, h4(e13))) | ~ (h10(op1(e12, e11)) = op2(e25, h9(e13))) | ~ (e20 = h10(op1(e12, e12))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), e24)) | ~ (h10(op1(e12, e15)) = op2(e25, h30(e13))) | ~ (h10(op1(e13, e10)) = op2(e25, h1(e13))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), e21)) | ~ (e24 = h10(op1(e13, e13))) | ~ (h10(op1(e13, e14)) = op2(e25, h22(e13))) | ~ (h10(op1(e13, e15)) = op2(e25, h27(e13))) | ~ (h4(e13) = h10(op1(e14, e10))) | ~ (h9(e13) = h10(op1(e14, e11))) | ~ (h10(op1(e14, e12)) = op2(e24, h10(e12))) | ~ (h10(op1(e14, e13)) = op2(e24, h10(e13))) | ~ (op2(e24, e24) = h10(op1(e14, e14))) | ~ (h30(e13) = h10(op1(e14, e15))) | ~ (h5(e13) = h10(op1(e15, e10))) | ~ (h10(e13) = h10(op1(e15, e11))) | ~ (h10(op1(e15, e12)) = op2(e25, h10(e12))) | ~ (h10(op1(e15, e13)) = op2(e25, h10(e13))) | ~ (h10(e12) = h10(op1(e15, e14))) | ~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e15)) = op2(h10(e10), h10(e15))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(forward_demodulation, [], [f15109, f1536])).
fof(f15109, plain, (~ (h22(e13) = h10(op1(e11, e14))) | ~ (h27(e13) = h10(op1(e11, e15))) | ~ (h10(op1(e12, e10)) = op2(e25, h4(e13))) | ~ (h10(op1(e12, e11)) = op2(e25, h9(e13))) | ~ (e20 = h10(op1(e12, e12))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), e24)) | ~ (h10(op1(e12, e15)) = op2(e25, h30(e13))) | ~ (h10(op1(e13, e10)) = op2(e25, h1(e13))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), e21)) | ~ (e24 = h10(op1(e13, e13))) | ~ (h10(op1(e13, e14)) = op2(e25, h22(e13))) | ~ (h10(op1(e13, e15)) = op2(e25, h27(e13))) | ~ (h4(e13) = h10(op1(e14, e10))) | ~ (h9(e13) = h10(op1(e14, e11))) | ~ (h10(op1(e14, e12)) = op2(e24, h10(e12))) | ~ (h10(op1(e14, e13)) = op2(e24, h10(e13))) | ~ (op2(e24, e24) = h10(op1(e14, e14))) | ~ (h30(e13) = h10(op1(e14, e15))) | ~ (h5(e13) = h10(op1(e15, e10))) | ~ (h10(e13) = h10(op1(e15, e11))) | ~ (h10(op1(e15, e12)) = op2(e25, h10(e12))) | ~ (h10(op1(e15, e13)) = op2(e25, h10(e13))) | ~ (h10(e12) = h10(op1(e15, e14))) | ~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e15)) = op2(h10(e10), h10(e15))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(forward_demodulation, [], [f15108, f1612])).
fof(f15108, plain, (~ (op2(e21, e24) = h10(op1(e11, e14))) | ~ (h27(e13) = h10(op1(e11, e15))) | ~ (h10(op1(e12, e10)) = op2(e25, h4(e13))) | ~ (h10(op1(e12, e11)) = op2(e25, h9(e13))) | ~ (e20 = h10(op1(e12, e12))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), e24)) | ~ (h10(op1(e12, e15)) = op2(e25, h30(e13))) | ~ (h10(op1(e13, e10)) = op2(e25, h1(e13))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), e21)) | ~ (e24 = h10(op1(e13, e13))) | ~ (h10(op1(e13, e14)) = op2(e25, h22(e13))) | ~ (h10(op1(e13, e15)) = op2(e25, h27(e13))) | ~ (h4(e13) = h10(op1(e14, e10))) | ~ (h9(e13) = h10(op1(e14, e11))) | ~ (h10(op1(e14, e12)) = op2(e24, h10(e12))) | ~ (h10(op1(e14, e13)) = op2(e24, h10(e13))) | ~ (op2(e24, e24) = h10(op1(e14, e14))) | ~ (h30(e13) = h10(op1(e14, e15))) | ~ (h5(e13) = h10(op1(e15, e10))) | ~ (h10(e13) = h10(op1(e15, e11))) | ~ (h10(op1(e15, e12)) = op2(e25, h10(e12))) | ~ (h10(op1(e15, e13)) = op2(e25, h10(e13))) | ~ (h10(e12) = h10(op1(e15, e14))) | ~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e15)) = op2(h10(e10), h10(e15))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(forward_demodulation, [], [f15107, f1536])).
fof(f15107, plain, (~ (h10(op1(e11, e14)) = op2(h10(e11), e24)) | ~ (h27(e13) = h10(op1(e11, e15))) | ~ (h10(op1(e12, e10)) = op2(e25, h4(e13))) | ~ (h10(op1(e12, e11)) = op2(e25, h9(e13))) | ~ (e20 = h10(op1(e12, e12))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), e24)) | ~ (h10(op1(e12, e15)) = op2(e25, h30(e13))) | ~ (h10(op1(e13, e10)) = op2(e25, h1(e13))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), e21)) | ~ (e24 = h10(op1(e13, e13))) | ~ (h10(op1(e13, e14)) = op2(e25, h22(e13))) | ~ (h10(op1(e13, e15)) = op2(e25, h27(e13))) | ~ (h4(e13) = h10(op1(e14, e10))) | ~ (h9(e13) = h10(op1(e14, e11))) | ~ (h10(op1(e14, e12)) = op2(e24, h10(e12))) | ~ (h10(op1(e14, e13)) = op2(e24, h10(e13))) | ~ (op2(e24, e24) = h10(op1(e14, e14))) | ~ (h30(e13) = h10(op1(e14, e15))) | ~ (h5(e13) = h10(op1(e15, e10))) | ~ (h10(e13) = h10(op1(e15, e11))) | ~ (h10(op1(e15, e12)) = op2(e25, h10(e12))) | ~ (h10(op1(e15, e13)) = op2(e25, h10(e13))) | ~ (h10(e12) = h10(op1(e15, e14))) | ~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e15)) = op2(h10(e10), h10(e15))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(forward_demodulation, [], [f15106, f5319])).
fof(f15106, plain, (~ (h27(e13) = h10(op1(e11, e15))) | ~ (h10(op1(e12, e10)) = op2(e25, h4(e13))) | ~ (h10(op1(e12, e11)) = op2(e25, h9(e13))) | ~ (e20 = h10(op1(e12, e12))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), e24)) | ~ (h10(op1(e12, e15)) = op2(e25, h30(e13))) | ~ (h10(op1(e13, e10)) = op2(e25, h1(e13))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), e21)) | ~ (e24 = h10(op1(e13, e13))) | ~ (h10(op1(e13, e14)) = op2(e25, h22(e13))) | ~ (h10(op1(e13, e15)) = op2(e25, h27(e13))) | ~ (h4(e13) = h10(op1(e14, e10))) | ~ (h9(e13) = h10(op1(e14, e11))) | ~ (h10(op1(e14, e12)) = op2(e24, h10(e12))) | ~ (h10(op1(e14, e13)) = op2(e24, h10(e13))) | ~ (op2(e24, e24) = h10(op1(e14, e14))) | ~ (h30(e13) = h10(op1(e14, e15))) | ~ (h5(e13) = h10(op1(e15, e10))) | ~ (h10(e13) = h10(op1(e15, e11))) | ~ (h10(op1(e15, e12)) = op2(e25, h10(e12))) | ~ (h10(op1(e15, e13)) = op2(e25, h10(e13))) | ~ (h10(e12) = h10(op1(e15, e14))) | ~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e11, e14)) = op2(h10(e11), h10(e14))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e15)) = op2(h10(e10), h10(e15))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(forward_demodulation, [], [f15105, f1642])).
fof(f15105, plain, (~ (op2(e21, e25) = h10(op1(e11, e15))) | ~ (h10(op1(e12, e10)) = op2(e25, h4(e13))) | ~ (h10(op1(e12, e11)) = op2(e25, h9(e13))) | ~ (e20 = h10(op1(e12, e12))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), e24)) | ~ (h10(op1(e12, e15)) = op2(e25, h30(e13))) | ~ (h10(op1(e13, e10)) = op2(e25, h1(e13))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), e21)) | ~ (e24 = h10(op1(e13, e13))) | ~ (h10(op1(e13, e14)) = op2(e25, h22(e13))) | ~ (h10(op1(e13, e15)) = op2(e25, h27(e13))) | ~ (h4(e13) = h10(op1(e14, e10))) | ~ (h9(e13) = h10(op1(e14, e11))) | ~ (h10(op1(e14, e12)) = op2(e24, h10(e12))) | ~ (h10(op1(e14, e13)) = op2(e24, h10(e13))) | ~ (op2(e24, e24) = h10(op1(e14, e14))) | ~ (h30(e13) = h10(op1(e14, e15))) | ~ (h5(e13) = h10(op1(e15, e10))) | ~ (h10(e13) = h10(op1(e15, e11))) | ~ (h10(op1(e15, e12)) = op2(e25, h10(e12))) | ~ (h10(op1(e15, e13)) = op2(e25, h10(e13))) | ~ (h10(e12) = h10(op1(e15, e14))) | ~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e11, e14)) = op2(h10(e11), h10(e14))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e15)) = op2(h10(e10), h10(e15))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(forward_demodulation, [], [f15104, f1536])).
fof(f15104, plain, (~ (h10(op1(e11, e15)) = op2(h10(e11), e25)) | ~ (h10(op1(e12, e10)) = op2(e25, h4(e13))) | ~ (h10(op1(e12, e11)) = op2(e25, h9(e13))) | ~ (e20 = h10(op1(e12, e12))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), e24)) | ~ (h10(op1(e12, e15)) = op2(e25, h30(e13))) | ~ (h10(op1(e13, e10)) = op2(e25, h1(e13))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), e21)) | ~ (e24 = h10(op1(e13, e13))) | ~ (h10(op1(e13, e14)) = op2(e25, h22(e13))) | ~ (h10(op1(e13, e15)) = op2(e25, h27(e13))) | ~ (h4(e13) = h10(op1(e14, e10))) | ~ (h9(e13) = h10(op1(e14, e11))) | ~ (h10(op1(e14, e12)) = op2(e24, h10(e12))) | ~ (h10(op1(e14, e13)) = op2(e24, h10(e13))) | ~ (op2(e24, e24) = h10(op1(e14, e14))) | ~ (h30(e13) = h10(op1(e14, e15))) | ~ (h5(e13) = h10(op1(e15, e10))) | ~ (h10(e13) = h10(op1(e15, e11))) | ~ (h10(op1(e15, e12)) = op2(e25, h10(e12))) | ~ (h10(op1(e15, e13)) = op2(e25, h10(e13))) | ~ (h10(e12) = h10(op1(e15, e14))) | ~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e11, e14)) = op2(h10(e11), h10(e14))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e15)) = op2(h10(e10), h10(e15))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(forward_demodulation, [], [f15103, f1537])).
fof(f15103, plain, (~ (h10(op1(e12, e10)) = op2(e25, h4(e13))) | ~ (h10(op1(e12, e11)) = op2(e25, h9(e13))) | ~ (e20 = h10(op1(e12, e12))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), e24)) | ~ (h10(op1(e12, e15)) = op2(e25, h30(e13))) | ~ (h10(op1(e13, e10)) = op2(e25, h1(e13))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), e21)) | ~ (e24 = h10(op1(e13, e13))) | ~ (h10(op1(e13, e14)) = op2(e25, h22(e13))) | ~ (h10(op1(e13, e15)) = op2(e25, h27(e13))) | ~ (h4(e13) = h10(op1(e14, e10))) | ~ (h9(e13) = h10(op1(e14, e11))) | ~ (h10(op1(e14, e12)) = op2(e24, h10(e12))) | ~ (h10(op1(e14, e13)) = op2(e24, h10(e13))) | ~ (op2(e24, e24) = h10(op1(e14, e14))) | ~ (h30(e13) = h10(op1(e14, e15))) | ~ (h5(e13) = h10(op1(e15, e10))) | ~ (h10(e13) = h10(op1(e15, e11))) | ~ (h10(op1(e15, e12)) = op2(e25, h10(e12))) | ~ (h10(op1(e15, e13)) = op2(e25, h10(e13))) | ~ (h10(e12) = h10(op1(e15, e14))) | ~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e11, e15)) = op2(h10(e11), h10(e15))) | ~ (h10(op1(e11, e14)) = op2(h10(e11), h10(e14))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e15)) = op2(h10(e10), h10(e15))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(forward_demodulation, [], [f15102, f5359])).
fof(f5359, plain, (op2(e25, h4(e13)) = op2(h10(e12), e20)), inference(backward_demodulation, [], [f5178, f5339])).
fof(f5178, plain, (op2(op2(e25, e24), e20) = op2(e25, h4(e13))), inference(backward_demodulation, [], [f887, f1504])).
fof(f887, plain, (op2(op2(e25, e24), e20) = op2(e25, op2(e24, e20))), inference(cnf_transformation, [], [f6])).
fof(f15102, plain, (~ (h10(op1(e12, e10)) = op2(h10(e12), e20)) | ~ (h10(op1(e12, e11)) = op2(e25, h9(e13))) | ~ (e20 = h10(op1(e12, e12))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), e24)) | ~ (h10(op1(e12, e15)) = op2(e25, h30(e13))) | ~ (h10(op1(e13, e10)) = op2(e25, h1(e13))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), e21)) | ~ (e24 = h10(op1(e13, e13))) | ~ (h10(op1(e13, e14)) = op2(e25, h22(e13))) | ~ (h10(op1(e13, e15)) = op2(e25, h27(e13))) | ~ (h4(e13) = h10(op1(e14, e10))) | ~ (h9(e13) = h10(op1(e14, e11))) | ~ (h10(op1(e14, e12)) = op2(e24, h10(e12))) | ~ (h10(op1(e14, e13)) = op2(e24, h10(e13))) | ~ (op2(e24, e24) = h10(op1(e14, e14))) | ~ (h30(e13) = h10(op1(e14, e15))) | ~ (h5(e13) = h10(op1(e15, e10))) | ~ (h10(e13) = h10(op1(e15, e11))) | ~ (h10(op1(e15, e12)) = op2(e25, h10(e12))) | ~ (h10(op1(e15, e13)) = op2(e25, h10(e13))) | ~ (h10(e12) = h10(op1(e15, e14))) | ~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e11, e15)) = op2(h10(e11), h10(e15))) | ~ (h10(op1(e11, e14)) = op2(h10(e11), h10(e14))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e15)) = op2(h10(e10), h10(e15))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(forward_demodulation, [], [f15101, f5365])).
fof(f15101, plain, (~ (h10(op1(e12, e11)) = op2(e25, h9(e13))) | ~ (e20 = h10(op1(e12, e12))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), e24)) | ~ (h10(op1(e12, e15)) = op2(e25, h30(e13))) | ~ (h10(op1(e13, e10)) = op2(e25, h1(e13))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), e21)) | ~ (e24 = h10(op1(e13, e13))) | ~ (h10(op1(e13, e14)) = op2(e25, h22(e13))) | ~ (h10(op1(e13, e15)) = op2(e25, h27(e13))) | ~ (h4(e13) = h10(op1(e14, e10))) | ~ (h9(e13) = h10(op1(e14, e11))) | ~ (h10(op1(e14, e12)) = op2(e24, h10(e12))) | ~ (h10(op1(e14, e13)) = op2(e24, h10(e13))) | ~ (op2(e24, e24) = h10(op1(e14, e14))) | ~ (h30(e13) = h10(op1(e14, e15))) | ~ (h5(e13) = h10(op1(e15, e10))) | ~ (h10(e13) = h10(op1(e15, e11))) | ~ (h10(op1(e15, e12)) = op2(e25, h10(e12))) | ~ (h10(op1(e15, e13)) = op2(e25, h10(e13))) | ~ (h10(e12) = h10(op1(e15, e14))) | ~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) | ~ (h10(op1(e11, e15)) = op2(h10(e11), h10(e15))) | ~ (h10(op1(e11, e14)) = op2(h10(e11), h10(e14))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e15)) = op2(h10(e10), h10(e15))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(forward_demodulation, [], [f15100, f5360])).
fof(f5360, plain, (op2(e25, h9(e13)) = op2(h10(e12), e21)), inference(backward_demodulation, [], [f5305, f5339])).
fof(f5305, plain, (op2(op2(e25, e24), e21) = op2(e25, h9(e13))), inference(backward_demodulation, [], [f888, f1534])).
fof(f888, plain, (op2(op2(e25, e24), e21) = op2(e25, op2(e24, e21))), inference(cnf_transformation, [], [f6])).
fof(f15100, plain, (~ (h10(op1(e12, e11)) = op2(h10(e12), e21)) | ~ (e20 = h10(op1(e12, e12))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), e24)) | ~ (h10(op1(e12, e15)) = op2(e25, h30(e13))) | ~ (h10(op1(e13, e10)) = op2(e25, h1(e13))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), e21)) | ~ (e24 = h10(op1(e13, e13))) | ~ (h10(op1(e13, e14)) = op2(e25, h22(e13))) | ~ (h10(op1(e13, e15)) = op2(e25, h27(e13))) | ~ (h4(e13) = h10(op1(e14, e10))) | ~ (h9(e13) = h10(op1(e14, e11))) | ~ (h10(op1(e14, e12)) = op2(e24, h10(e12))) | ~ (h10(op1(e14, e13)) = op2(e24, h10(e13))) | ~ (op2(e24, e24) = h10(op1(e14, e14))) | ~ (h30(e13) = h10(op1(e14, e15))) | ~ (h5(e13) = h10(op1(e15, e10))) | ~ (h10(e13) = h10(op1(e15, e11))) | ~ (h10(op1(e15, e12)) = op2(e25, h10(e12))) | ~ (h10(op1(e15, e13)) = op2(e25, h10(e13))) | ~ (h10(e12) = h10(op1(e15, e14))) | ~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) | ~ (h10(op1(e11, e15)) = op2(h10(e11), h10(e15))) | ~ (h10(op1(e11, e14)) = op2(h10(e11), h10(e14))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e15)) = op2(h10(e10), h10(e15))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(forward_demodulation, [], [f15099, f1536])).
fof(f15099, plain, (~ (e20 = h10(op1(e12, e12))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), e24)) | ~ (h10(op1(e12, e15)) = op2(e25, h30(e13))) | ~ (h10(op1(e13, e10)) = op2(e25, h1(e13))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), e21)) | ~ (e24 = h10(op1(e13, e13))) | ~ (h10(op1(e13, e14)) = op2(e25, h22(e13))) | ~ (h10(op1(e13, e15)) = op2(e25, h27(e13))) | ~ (h4(e13) = h10(op1(e14, e10))) | ~ (h9(e13) = h10(op1(e14, e11))) | ~ (h10(op1(e14, e12)) = op2(e24, h10(e12))) | ~ (h10(op1(e14, e13)) = op2(e24, h10(e13))) | ~ (op2(e24, e24) = h10(op1(e14, e14))) | ~ (h30(e13) = h10(op1(e14, e15))) | ~ (h5(e13) = h10(op1(e15, e10))) | ~ (h10(e13) = h10(op1(e15, e11))) | ~ (h10(op1(e15, e12)) = op2(e25, h10(e12))) | ~ (h10(op1(e15, e13)) = op2(e25, h10(e13))) | ~ (h10(e12) = h10(op1(e15, e14))) | ~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) | ~ (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) | ~ (h10(op1(e11, e15)) = op2(h10(e11), h10(e15))) | ~ (h10(op1(e11, e14)) = op2(h10(e11), h10(e14))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e15)) = op2(h10(e10), h10(e15))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(forward_demodulation, [], [f15098, f5358])).
fof(f15098, plain, (~ (h10(op1(e12, e14)) = op2(h10(e12), e24)) | ~ (h10(op1(e12, e15)) = op2(e25, h30(e13))) | ~ (h10(op1(e13, e10)) = op2(e25, h1(e13))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), e21)) | ~ (e24 = h10(op1(e13, e13))) | ~ (h10(op1(e13, e14)) = op2(e25, h22(e13))) | ~ (h10(op1(e13, e15)) = op2(e25, h27(e13))) | ~ (h4(e13) = h10(op1(e14, e10))) | ~ (h9(e13) = h10(op1(e14, e11))) | ~ (h10(op1(e14, e12)) = op2(e24, h10(e12))) | ~ (h10(op1(e14, e13)) = op2(e24, h10(e13))) | ~ (op2(e24, e24) = h10(op1(e14, e14))) | ~ (h30(e13) = h10(op1(e14, e15))) | ~ (h5(e13) = h10(op1(e15, e10))) | ~ (h10(e13) = h10(op1(e15, e11))) | ~ (h10(op1(e15, e12)) = op2(e25, h10(e12))) | ~ (h10(op1(e15, e13)) = op2(e25, h10(e13))) | ~ (h10(e12) = h10(op1(e15, e14))) | ~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) | ~ (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) | ~ (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) | ~ (h10(op1(e11, e15)) = op2(h10(e11), h10(e15))) | ~ (h10(op1(e11, e14)) = op2(h10(e11), h10(e14))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e15)) = op2(h10(e10), h10(e15))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(forward_demodulation, [], [f15097, f5319])).
fof(f15097, plain, (~ (h10(op1(e12, e15)) = op2(e25, h30(e13))) | ~ (h10(op1(e13, e10)) = op2(e25, h1(e13))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), e21)) | ~ (e24 = h10(op1(e13, e13))) | ~ (h10(op1(e13, e14)) = op2(e25, h22(e13))) | ~ (h10(op1(e13, e15)) = op2(e25, h27(e13))) | ~ (h4(e13) = h10(op1(e14, e10))) | ~ (h9(e13) = h10(op1(e14, e11))) | ~ (h10(op1(e14, e12)) = op2(e24, h10(e12))) | ~ (h10(op1(e14, e13)) = op2(e24, h10(e13))) | ~ (op2(e24, e24) = h10(op1(e14, e14))) | ~ (h30(e13) = h10(op1(e14, e15))) | ~ (h5(e13) = h10(op1(e15, e10))) | ~ (h10(e13) = h10(op1(e15, e11))) | ~ (h10(op1(e15, e12)) = op2(e25, h10(e12))) | ~ (h10(op1(e15, e13)) = op2(e25, h10(e13))) | ~ (h10(e12) = h10(op1(e15, e14))) | ~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), h10(e14))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) | ~ (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) | ~ (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) | ~ (h10(op1(e11, e15)) = op2(h10(e11), h10(e15))) | ~ (h10(op1(e11, e14)) = op2(h10(e11), h10(e14))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e15)) = op2(h10(e10), h10(e15))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(forward_demodulation, [], [f15096, f5802])).
fof(f15096, plain, (~ (h10(op1(e12, e15)) = op2(h10(e12), e25)) | ~ (h10(op1(e13, e10)) = op2(e25, h1(e13))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), e21)) | ~ (e24 = h10(op1(e13, e13))) | ~ (h10(op1(e13, e14)) = op2(e25, h22(e13))) | ~ (h10(op1(e13, e15)) = op2(e25, h27(e13))) | ~ (h4(e13) = h10(op1(e14, e10))) | ~ (h9(e13) = h10(op1(e14, e11))) | ~ (h10(op1(e14, e12)) = op2(e24, h10(e12))) | ~ (h10(op1(e14, e13)) = op2(e24, h10(e13))) | ~ (op2(e24, e24) = h10(op1(e14, e14))) | ~ (h30(e13) = h10(op1(e14, e15))) | ~ (h5(e13) = h10(op1(e15, e10))) | ~ (h10(e13) = h10(op1(e15, e11))) | ~ (h10(op1(e15, e12)) = op2(e25, h10(e12))) | ~ (h10(op1(e15, e13)) = op2(e25, h10(e13))) | ~ (h10(e12) = h10(op1(e15, e14))) | ~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), h10(e14))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) | ~ (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) | ~ (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) | ~ (h10(op1(e11, e15)) = op2(h10(e11), h10(e15))) | ~ (h10(op1(e11, e14)) = op2(h10(e11), h10(e14))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e15)) = op2(h10(e10), h10(e15))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(forward_demodulation, [], [f15095, f1537])).
fof(f15095, plain, (~ (h10(op1(e13, e10)) = op2(e25, h1(e13))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), e21)) | ~ (e24 = h10(op1(e13, e13))) | ~ (h10(op1(e13, e14)) = op2(e25, h22(e13))) | ~ (h10(op1(e13, e15)) = op2(e25, h27(e13))) | ~ (h4(e13) = h10(op1(e14, e10))) | ~ (h9(e13) = h10(op1(e14, e11))) | ~ (h10(op1(e14, e12)) = op2(e24, h10(e12))) | ~ (h10(op1(e14, e13)) = op2(e24, h10(e13))) | ~ (op2(e24, e24) = h10(op1(e14, e14))) | ~ (h30(e13) = h10(op1(e14, e15))) | ~ (h5(e13) = h10(op1(e15, e10))) | ~ (h10(e13) = h10(op1(e15, e11))) | ~ (h10(op1(e15, e12)) = op2(e25, h10(e12))) | ~ (h10(op1(e15, e13)) = op2(e25, h10(e13))) | ~ (h10(e12) = h10(op1(e15, e14))) | ~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e12, e15)) = op2(h10(e12), h10(e15))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), h10(e14))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) | ~ (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) | ~ (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) | ~ (h10(op1(e11, e15)) = op2(h10(e11), h10(e15))) | ~ (h10(op1(e11, e14)) = op2(h10(e11), h10(e14))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e15)) = op2(h10(e10), h10(e15))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(forward_demodulation, [], [f15094, f5337])).
fof(f5337, plain, (op2(e25, h1(e13)) = op2(h10(e13), e20)), inference(backward_demodulation, [], [f5097, f1540])).
fof(f5097, plain, (op2(op2(e25, e21), e20) = op2(e25, h1(e13))), inference(backward_demodulation, [], [f869, f1486])).
fof(f869, plain, (op2(op2(e25, e21), e20) = op2(e25, op2(e21, e20))), inference(cnf_transformation, [], [f6])).
fof(f15094, plain, (~ (h10(op1(e13, e10)) = op2(h10(e13), e20)) | ~ (h10(op1(e13, e11)) = op2(h10(e13), e21)) | ~ (e24 = h10(op1(e13, e13))) | ~ (h10(op1(e13, e14)) = op2(e25, h22(e13))) | ~ (h10(op1(e13, e15)) = op2(e25, h27(e13))) | ~ (h4(e13) = h10(op1(e14, e10))) | ~ (h9(e13) = h10(op1(e14, e11))) | ~ (h10(op1(e14, e12)) = op2(e24, h10(e12))) | ~ (h10(op1(e14, e13)) = op2(e24, h10(e13))) | ~ (op2(e24, e24) = h10(op1(e14, e14))) | ~ (h30(e13) = h10(op1(e14, e15))) | ~ (h5(e13) = h10(op1(e15, e10))) | ~ (h10(e13) = h10(op1(e15, e11))) | ~ (h10(op1(e15, e12)) = op2(e25, h10(e12))) | ~ (h10(op1(e15, e13)) = op2(e25, h10(e13))) | ~ (h10(e12) = h10(op1(e15, e14))) | ~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e12, e15)) = op2(h10(e12), h10(e15))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), h10(e14))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) | ~ (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) | ~ (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) | ~ (h10(op1(e11, e15)) = op2(h10(e11), h10(e15))) | ~ (h10(op1(e11, e14)) = op2(h10(e11), h10(e14))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e15)) = op2(h10(e10), h10(e15))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(forward_demodulation, [], [f15093, f5365])).
fof(f15093, plain, (~ (h10(op1(e13, e11)) = op2(h10(e13), e21)) | ~ (e24 = h10(op1(e13, e13))) | ~ (h10(op1(e13, e14)) = op2(e25, h22(e13))) | ~ (h10(op1(e13, e15)) = op2(e25, h27(e13))) | ~ (h4(e13) = h10(op1(e14, e10))) | ~ (h9(e13) = h10(op1(e14, e11))) | ~ (h10(op1(e14, e12)) = op2(e24, h10(e12))) | ~ (h10(op1(e14, e13)) = op2(e24, h10(e13))) | ~ (op2(e24, e24) = h10(op1(e14, e14))) | ~ (h30(e13) = h10(op1(e14, e15))) | ~ (h5(e13) = h10(op1(e15, e10))) | ~ (h10(e13) = h10(op1(e15, e11))) | ~ (h10(op1(e15, e12)) = op2(e25, h10(e12))) | ~ (h10(op1(e15, e13)) = op2(e25, h10(e13))) | ~ (h10(e12) = h10(op1(e15, e14))) | ~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) | ~ (h10(op1(e12, e15)) = op2(h10(e12), h10(e15))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), h10(e14))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) | ~ (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) | ~ (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) | ~ (h10(op1(e11, e15)) = op2(h10(e11), h10(e15))) | ~ (h10(op1(e11, e14)) = op2(h10(e11), h10(e14))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e15)) = op2(h10(e10), h10(e15))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(forward_demodulation, [], [f15092, f1536])).
fof(f15092, plain, (~ (e24 = h10(op1(e13, e13))) | ~ (h10(op1(e13, e14)) = op2(e25, h22(e13))) | ~ (h10(op1(e13, e15)) = op2(e25, h27(e13))) | ~ (h4(e13) = h10(op1(e14, e10))) | ~ (h9(e13) = h10(op1(e14, e11))) | ~ (h10(op1(e14, e12)) = op2(e24, h10(e12))) | ~ (h10(op1(e14, e13)) = op2(e24, h10(e13))) | ~ (op2(e24, e24) = h10(op1(e14, e14))) | ~ (h30(e13) = h10(op1(e14, e15))) | ~ (h5(e13) = h10(op1(e15, e10))) | ~ (h10(e13) = h10(op1(e15, e11))) | ~ (h10(op1(e15, e12)) = op2(e25, h10(e12))) | ~ (h10(op1(e15, e13)) = op2(e25, h10(e13))) | ~ (h10(e12) = h10(op1(e15, e14))) | ~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) | ~ (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) | ~ (h10(op1(e12, e15)) = op2(h10(e12), h10(e15))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), h10(e14))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) | ~ (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) | ~ (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) | ~ (h10(op1(e11, e15)) = op2(h10(e11), h10(e15))) | ~ (h10(op1(e11, e14)) = op2(h10(e11), h10(e14))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e15)) = op2(h10(e10), h10(e15))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(forward_demodulation, [], [f15091, f5336])).
fof(f15091, plain, (~ (h10(op1(e13, e14)) = op2(e25, h22(e13))) | ~ (h10(op1(e13, e15)) = op2(e25, h27(e13))) | ~ (h4(e13) = h10(op1(e14, e10))) | ~ (h9(e13) = h10(op1(e14, e11))) | ~ (h10(op1(e14, e12)) = op2(e24, h10(e12))) | ~ (h10(op1(e14, e13)) = op2(e24, h10(e13))) | ~ (op2(e24, e24) = h10(op1(e14, e14))) | ~ (h30(e13) = h10(op1(e14, e15))) | ~ (h5(e13) = h10(op1(e15, e10))) | ~ (h10(e13) = h10(op1(e15, e11))) | ~ (h10(op1(e15, e12)) = op2(e25, h10(e12))) | ~ (h10(op1(e15, e13)) = op2(e25, h10(e13))) | ~ (h10(e12) = h10(op1(e15, e14))) | ~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e13, e13)) = op2(h10(e13), h10(e13))) | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) | ~ (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) | ~ (h10(op1(e12, e15)) = op2(h10(e12), h10(e15))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), h10(e14))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) | ~ (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) | ~ (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) | ~ (h10(op1(e11, e15)) = op2(h10(e11), h10(e15))) | ~ (h10(op1(e11, e14)) = op2(h10(e11), h10(e14))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e15)) = op2(h10(e10), h10(e15))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(forward_demodulation, [], [f15090, f5643])).
fof(f15090, plain, (~ (h10(op1(e13, e14)) = op2(h10(e13), e24)) | ~ (h10(op1(e13, e15)) = op2(e25, h27(e13))) | ~ (h4(e13) = h10(op1(e14, e10))) | ~ (h9(e13) = h10(op1(e14, e11))) | ~ (h10(op1(e14, e12)) = op2(e24, h10(e12))) | ~ (h10(op1(e14, e13)) = op2(e24, h10(e13))) | ~ (op2(e24, e24) = h10(op1(e14, e14))) | ~ (h30(e13) = h10(op1(e14, e15))) | ~ (h5(e13) = h10(op1(e15, e10))) | ~ (h10(e13) = h10(op1(e15, e11))) | ~ (h10(op1(e15, e12)) = op2(e25, h10(e12))) | ~ (h10(op1(e15, e13)) = op2(e25, h10(e13))) | ~ (h10(e12) = h10(op1(e15, e14))) | ~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e13, e13)) = op2(h10(e13), h10(e13))) | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) | ~ (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) | ~ (h10(op1(e12, e15)) = op2(h10(e12), h10(e15))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), h10(e14))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) | ~ (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) | ~ (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) | ~ (h10(op1(e11, e15)) = op2(h10(e11), h10(e15))) | ~ (h10(op1(e11, e14)) = op2(h10(e11), h10(e14))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e15)) = op2(h10(e10), h10(e15))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(forward_demodulation, [], [f15089, f5319])).
fof(f15089, plain, (~ (h10(op1(e13, e15)) = op2(e25, h27(e13))) | ~ (h4(e13) = h10(op1(e14, e10))) | ~ (h9(e13) = h10(op1(e14, e11))) | ~ (h10(op1(e14, e12)) = op2(e24, h10(e12))) | ~ (h10(op1(e14, e13)) = op2(e24, h10(e13))) | ~ (op2(e24, e24) = h10(op1(e14, e14))) | ~ (h30(e13) = h10(op1(e14, e15))) | ~ (h5(e13) = h10(op1(e15, e10))) | ~ (h10(e13) = h10(op1(e15, e11))) | ~ (h10(op1(e15, e12)) = op2(e25, h10(e12))) | ~ (h10(op1(e15, e13)) = op2(e25, h10(e13))) | ~ (h10(e12) = h10(op1(e15, e14))) | ~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e13, e14)) = op2(h10(e13), h10(e14))) | ~ (h10(op1(e13, e13)) = op2(h10(e13), h10(e13))) | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) | ~ (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) | ~ (h10(op1(e12, e15)) = op2(h10(e12), h10(e15))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), h10(e14))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) | ~ (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) | ~ (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) | ~ (h10(op1(e11, e15)) = op2(h10(e11), h10(e15))) | ~ (h10(op1(e11, e14)) = op2(h10(e11), h10(e14))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e15)) = op2(h10(e10), h10(e15))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(forward_demodulation, [], [f15088, f5742])).
fof(f5742, plain, (op2(h10(e13), e25) = op2(e25, h27(e13))), inference(backward_demodulation, [], [f5329, f1642])).
fof(f5329, plain, (op2(e25, op2(e21, e25)) = op2(h10(e13), e25)), inference(backward_demodulation, [], [f874, f1540])).
fof(f874, plain, (op2(op2(e25, e21), e25) = op2(e25, op2(e21, e25))), inference(cnf_transformation, [], [f6])).
fof(f15088, plain, (~ (h10(op1(e13, e15)) = op2(h10(e13), e25)) | ~ (h4(e13) = h10(op1(e14, e10))) | ~ (h9(e13) = h10(op1(e14, e11))) | ~ (h10(op1(e14, e12)) = op2(e24, h10(e12))) | ~ (h10(op1(e14, e13)) = op2(e24, h10(e13))) | ~ (op2(e24, e24) = h10(op1(e14, e14))) | ~ (h30(e13) = h10(op1(e14, e15))) | ~ (h5(e13) = h10(op1(e15, e10))) | ~ (h10(e13) = h10(op1(e15, e11))) | ~ (h10(op1(e15, e12)) = op2(e25, h10(e12))) | ~ (h10(op1(e15, e13)) = op2(e25, h10(e13))) | ~ (h10(e12) = h10(op1(e15, e14))) | ~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e13, e14)) = op2(h10(e13), h10(e14))) | ~ (h10(op1(e13, e13)) = op2(h10(e13), h10(e13))) | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) | ~ (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) | ~ (h10(op1(e12, e15)) = op2(h10(e12), h10(e15))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), h10(e14))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) | ~ (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) | ~ (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) | ~ (h10(op1(e11, e15)) = op2(h10(e11), h10(e15))) | ~ (h10(op1(e11, e14)) = op2(h10(e11), h10(e14))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e15)) = op2(h10(e10), h10(e15))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(forward_demodulation, [], [f15087, f1537])).
fof(f15087, plain, (~ (h4(e13) = h10(op1(e14, e10))) | ~ (h9(e13) = h10(op1(e14, e11))) | ~ (h10(op1(e14, e12)) = op2(e24, h10(e12))) | ~ (h10(op1(e14, e13)) = op2(e24, h10(e13))) | ~ (op2(e24, e24) = h10(op1(e14, e14))) | ~ (h30(e13) = h10(op1(e14, e15))) | ~ (h5(e13) = h10(op1(e15, e10))) | ~ (h10(e13) = h10(op1(e15, e11))) | ~ (h10(op1(e15, e12)) = op2(e25, h10(e12))) | ~ (h10(op1(e15, e13)) = op2(e25, h10(e13))) | ~ (h10(e12) = h10(op1(e15, e14))) | ~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e13, e15)) = op2(h10(e13), h10(e15))) | ~ (h10(op1(e13, e14)) = op2(h10(e13), h10(e14))) | ~ (h10(op1(e13, e13)) = op2(h10(e13), h10(e13))) | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) | ~ (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) | ~ (h10(op1(e12, e15)) = op2(h10(e12), h10(e15))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), h10(e14))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) | ~ (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) | ~ (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) | ~ (h10(op1(e11, e15)) = op2(h10(e11), h10(e15))) | ~ (h10(op1(e11, e14)) = op2(h10(e11), h10(e14))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e15)) = op2(h10(e10), h10(e15))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(forward_demodulation, [], [f15086, f1504])).
fof(f15086, plain, (~ (op2(e24, e20) = h10(op1(e14, e10))) | ~ (h9(e13) = h10(op1(e14, e11))) | ~ (h10(op1(e14, e12)) = op2(e24, h10(e12))) | ~ (h10(op1(e14, e13)) = op2(e24, h10(e13))) | ~ (op2(e24, e24) = h10(op1(e14, e14))) | ~ (h30(e13) = h10(op1(e14, e15))) | ~ (h5(e13) = h10(op1(e15, e10))) | ~ (h10(e13) = h10(op1(e15, e11))) | ~ (h10(op1(e15, e12)) = op2(e25, h10(e12))) | ~ (h10(op1(e15, e13)) = op2(e25, h10(e13))) | ~ (h10(e12) = h10(op1(e15, e14))) | ~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e13, e15)) = op2(h10(e13), h10(e15))) | ~ (h10(op1(e13, e14)) = op2(h10(e13), h10(e14))) | ~ (h10(op1(e13, e13)) = op2(h10(e13), h10(e13))) | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) | ~ (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) | ~ (h10(op1(e12, e15)) = op2(h10(e12), h10(e15))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), h10(e14))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) | ~ (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) | ~ (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) | ~ (h10(op1(e11, e15)) = op2(h10(e11), h10(e15))) | ~ (h10(op1(e11, e14)) = op2(h10(e11), h10(e14))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e15)) = op2(h10(e10), h10(e15))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(forward_demodulation, [], [f15085, f5319])).
fof(f15085, plain, (~ (h10(op1(e14, e10)) = op2(h10(e14), e20)) | ~ (h9(e13) = h10(op1(e14, e11))) | ~ (h10(op1(e14, e12)) = op2(e24, h10(e12))) | ~ (h10(op1(e14, e13)) = op2(e24, h10(e13))) | ~ (op2(e24, e24) = h10(op1(e14, e14))) | ~ (h30(e13) = h10(op1(e14, e15))) | ~ (h5(e13) = h10(op1(e15, e10))) | ~ (h10(e13) = h10(op1(e15, e11))) | ~ (h10(op1(e15, e12)) = op2(e25, h10(e12))) | ~ (h10(op1(e15, e13)) = op2(e25, h10(e13))) | ~ (h10(e12) = h10(op1(e15, e14))) | ~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e13, e15)) = op2(h10(e13), h10(e15))) | ~ (h10(op1(e13, e14)) = op2(h10(e13), h10(e14))) | ~ (h10(op1(e13, e13)) = op2(h10(e13), h10(e13))) | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) | ~ (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) | ~ (h10(op1(e12, e15)) = op2(h10(e12), h10(e15))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), h10(e14))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) | ~ (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) | ~ (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) | ~ (h10(op1(e11, e15)) = op2(h10(e11), h10(e15))) | ~ (h10(op1(e11, e14)) = op2(h10(e11), h10(e14))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e15)) = op2(h10(e10), h10(e15))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(forward_demodulation, [], [f15084, f5365])).
fof(f15084, plain, (~ (h9(e13) = h10(op1(e14, e11))) | ~ (h10(op1(e14, e12)) = op2(e24, h10(e12))) | ~ (h10(op1(e14, e13)) = op2(e24, h10(e13))) | ~ (op2(e24, e24) = h10(op1(e14, e14))) | ~ (h30(e13) = h10(op1(e14, e15))) | ~ (h5(e13) = h10(op1(e15, e10))) | ~ (h10(e13) = h10(op1(e15, e11))) | ~ (h10(op1(e15, e12)) = op2(e25, h10(e12))) | ~ (h10(op1(e15, e13)) = op2(e25, h10(e13))) | ~ (h10(e12) = h10(op1(e15, e14))) | ~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e14, e10)) = op2(h10(e14), h10(e10))) | ~ (h10(op1(e13, e15)) = op2(h10(e13), h10(e15))) | ~ (h10(op1(e13, e14)) = op2(h10(e13), h10(e14))) | ~ (h10(op1(e13, e13)) = op2(h10(e13), h10(e13))) | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) | ~ (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) | ~ (h10(op1(e12, e15)) = op2(h10(e12), h10(e15))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), h10(e14))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) | ~ (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) | ~ (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) | ~ (h10(op1(e11, e15)) = op2(h10(e11), h10(e15))) | ~ (h10(op1(e11, e14)) = op2(h10(e11), h10(e14))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e15)) = op2(h10(e10), h10(e15))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(forward_demodulation, [], [f15083, f1534])).
fof(f15083, plain, (~ (op2(e24, e21) = h10(op1(e14, e11))) | ~ (h10(op1(e14, e12)) = op2(e24, h10(e12))) | ~ (h10(op1(e14, e13)) = op2(e24, h10(e13))) | ~ (op2(e24, e24) = h10(op1(e14, e14))) | ~ (h30(e13) = h10(op1(e14, e15))) | ~ (h5(e13) = h10(op1(e15, e10))) | ~ (h10(e13) = h10(op1(e15, e11))) | ~ (h10(op1(e15, e12)) = op2(e25, h10(e12))) | ~ (h10(op1(e15, e13)) = op2(e25, h10(e13))) | ~ (h10(e12) = h10(op1(e15, e14))) | ~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e14, e10)) = op2(h10(e14), h10(e10))) | ~ (h10(op1(e13, e15)) = op2(h10(e13), h10(e15))) | ~ (h10(op1(e13, e14)) = op2(h10(e13), h10(e14))) | ~ (h10(op1(e13, e13)) = op2(h10(e13), h10(e13))) | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) | ~ (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) | ~ (h10(op1(e12, e15)) = op2(h10(e12), h10(e15))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), h10(e14))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) | ~ (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) | ~ (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) | ~ (h10(op1(e11, e15)) = op2(h10(e11), h10(e15))) | ~ (h10(op1(e11, e14)) = op2(h10(e11), h10(e14))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e15)) = op2(h10(e10), h10(e15))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(forward_demodulation, [], [f15082, f5319])).
fof(f15082, plain, (~ (h10(op1(e14, e11)) = op2(h10(e14), e21)) | ~ (h10(op1(e14, e12)) = op2(e24, h10(e12))) | ~ (h10(op1(e14, e13)) = op2(e24, h10(e13))) | ~ (op2(e24, e24) = h10(op1(e14, e14))) | ~ (h30(e13) = h10(op1(e14, e15))) | ~ (h5(e13) = h10(op1(e15, e10))) | ~ (h10(e13) = h10(op1(e15, e11))) | ~ (h10(op1(e15, e12)) = op2(e25, h10(e12))) | ~ (h10(op1(e15, e13)) = op2(e25, h10(e13))) | ~ (h10(e12) = h10(op1(e15, e14))) | ~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e14, e10)) = op2(h10(e14), h10(e10))) | ~ (h10(op1(e13, e15)) = op2(h10(e13), h10(e15))) | ~ (h10(op1(e13, e14)) = op2(h10(e13), h10(e14))) | ~ (h10(op1(e13, e13)) = op2(h10(e13), h10(e13))) | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) | ~ (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) | ~ (h10(op1(e12, e15)) = op2(h10(e12), h10(e15))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), h10(e14))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) | ~ (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) | ~ (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) | ~ (h10(op1(e11, e15)) = op2(h10(e11), h10(e15))) | ~ (h10(op1(e11, e14)) = op2(h10(e11), h10(e14))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e15)) = op2(h10(e10), h10(e15))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(forward_demodulation, [], [f15081, f1536])).
fof(f15081, plain, (~ (h10(op1(e14, e12)) = op2(e24, h10(e12))) | ~ (h10(op1(e14, e13)) = op2(e24, h10(e13))) | ~ (op2(e24, e24) = h10(op1(e14, e14))) | ~ (h30(e13) = h10(op1(e14, e15))) | ~ (h5(e13) = h10(op1(e15, e10))) | ~ (h10(e13) = h10(op1(e15, e11))) | ~ (h10(op1(e15, e12)) = op2(e25, h10(e12))) | ~ (h10(op1(e15, e13)) = op2(e25, h10(e13))) | ~ (h10(e12) = h10(op1(e15, e14))) | ~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e14, e11)) = op2(h10(e14), h10(e11))) | ~ (h10(op1(e14, e10)) = op2(h10(e14), h10(e10))) | ~ (h10(op1(e13, e15)) = op2(h10(e13), h10(e15))) | ~ (h10(op1(e13, e14)) = op2(h10(e13), h10(e14))) | ~ (h10(op1(e13, e13)) = op2(h10(e13), h10(e13))) | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) | ~ (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) | ~ (h10(op1(e12, e15)) = op2(h10(e12), h10(e15))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), h10(e14))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) | ~ (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) | ~ (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) | ~ (h10(op1(e11, e15)) = op2(h10(e11), h10(e15))) | ~ (h10(op1(e11, e14)) = op2(h10(e11), h10(e14))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e15)) = op2(h10(e10), h10(e15))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(forward_demodulation, [], [f15080, f5319])).
fof(f15080, plain, (~ (h10(op1(e14, e13)) = op2(e24, h10(e13))) | ~ (op2(e24, e24) = h10(op1(e14, e14))) | ~ (h30(e13) = h10(op1(e14, e15))) | ~ (h5(e13) = h10(op1(e15, e10))) | ~ (h10(e13) = h10(op1(e15, e11))) | ~ (h10(op1(e15, e12)) = op2(e25, h10(e12))) | ~ (h10(op1(e15, e13)) = op2(e25, h10(e13))) | ~ (h10(e12) = h10(op1(e15, e14))) | ~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e14, e12)) = op2(h10(e14), h10(e12))) | ~ (h10(op1(e14, e11)) = op2(h10(e14), h10(e11))) | ~ (h10(op1(e14, e10)) = op2(h10(e14), h10(e10))) | ~ (h10(op1(e13, e15)) = op2(h10(e13), h10(e15))) | ~ (h10(op1(e13, e14)) = op2(h10(e13), h10(e14))) | ~ (h10(op1(e13, e13)) = op2(h10(e13), h10(e13))) | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) | ~ (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) | ~ (h10(op1(e12, e15)) = op2(h10(e12), h10(e15))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), h10(e14))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) | ~ (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) | ~ (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) | ~ (h10(op1(e11, e15)) = op2(h10(e11), h10(e15))) | ~ (h10(op1(e11, e14)) = op2(h10(e11), h10(e14))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e15)) = op2(h10(e10), h10(e15))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(forward_demodulation, [], [f15079, f5319])).
fof(f15079, plain, (~ (op2(e24, e24) = h10(op1(e14, e14))) | ~ (h30(e13) = h10(op1(e14, e15))) | ~ (h5(e13) = h10(op1(e15, e10))) | ~ (h10(e13) = h10(op1(e15, e11))) | ~ (h10(op1(e15, e12)) = op2(e25, h10(e12))) | ~ (h10(op1(e15, e13)) = op2(e25, h10(e13))) | ~ (h10(e12) = h10(op1(e15, e14))) | ~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e14, e13)) = op2(h10(e14), h10(e13))) | ~ (h10(op1(e14, e12)) = op2(h10(e14), h10(e12))) | ~ (h10(op1(e14, e11)) = op2(h10(e14), h10(e11))) | ~ (h10(op1(e14, e10)) = op2(h10(e14), h10(e10))) | ~ (h10(op1(e13, e15)) = op2(h10(e13), h10(e15))) | ~ (h10(op1(e13, e14)) = op2(h10(e13), h10(e14))) | ~ (h10(op1(e13, e13)) = op2(h10(e13), h10(e13))) | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) | ~ (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) | ~ (h10(op1(e12, e15)) = op2(h10(e12), h10(e15))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), h10(e14))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) | ~ (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) | ~ (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) | ~ (h10(op1(e11, e15)) = op2(h10(e11), h10(e15))) | ~ (h10(op1(e11, e14)) = op2(h10(e11), h10(e14))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e15)) = op2(h10(e10), h10(e15))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(forward_demodulation, [], [f15078, f5319])).
fof(f15078, plain, (~ (h30(e13) = h10(op1(e14, e15))) | ~ (h5(e13) = h10(op1(e15, e10))) | ~ (h10(e13) = h10(op1(e15, e11))) | ~ (h10(op1(e15, e12)) = op2(e25, h10(e12))) | ~ (h10(op1(e15, e13)) = op2(e25, h10(e13))) | ~ (h10(e12) = h10(op1(e15, e14))) | ~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e14, e14)) = op2(h10(e14), h10(e14))) | ~ (h10(op1(e14, e13)) = op2(h10(e14), h10(e13))) | ~ (h10(op1(e14, e12)) = op2(h10(e14), h10(e12))) | ~ (h10(op1(e14, e11)) = op2(h10(e14), h10(e11))) | ~ (h10(op1(e14, e10)) = op2(h10(e14), h10(e10))) | ~ (h10(op1(e13, e15)) = op2(h10(e13), h10(e15))) | ~ (h10(op1(e13, e14)) = op2(h10(e13), h10(e14))) | ~ (h10(op1(e13, e13)) = op2(h10(e13), h10(e13))) | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) | ~ (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) | ~ (h10(op1(e12, e15)) = op2(h10(e12), h10(e15))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), h10(e14))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) | ~ (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) | ~ (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) | ~ (h10(op1(e11, e15)) = op2(h10(e11), h10(e15))) | ~ (h10(op1(e11, e14)) = op2(h10(e11), h10(e14))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e15)) = op2(h10(e10), h10(e15))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(forward_demodulation, [], [f15077, f1660])).
fof(f15077, plain, (~ (op2(e24, e25) = h10(op1(e14, e15))) | ~ (h5(e13) = h10(op1(e15, e10))) | ~ (h10(e13) = h10(op1(e15, e11))) | ~ (h10(op1(e15, e12)) = op2(e25, h10(e12))) | ~ (h10(op1(e15, e13)) = op2(e25, h10(e13))) | ~ (h10(e12) = h10(op1(e15, e14))) | ~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e14, e14)) = op2(h10(e14), h10(e14))) | ~ (h10(op1(e14, e13)) = op2(h10(e14), h10(e13))) | ~ (h10(op1(e14, e12)) = op2(h10(e14), h10(e12))) | ~ (h10(op1(e14, e11)) = op2(h10(e14), h10(e11))) | ~ (h10(op1(e14, e10)) = op2(h10(e14), h10(e10))) | ~ (h10(op1(e13, e15)) = op2(h10(e13), h10(e15))) | ~ (h10(op1(e13, e14)) = op2(h10(e13), h10(e14))) | ~ (h10(op1(e13, e13)) = op2(h10(e13), h10(e13))) | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) | ~ (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) | ~ (h10(op1(e12, e15)) = op2(h10(e12), h10(e15))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), h10(e14))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) | ~ (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) | ~ (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) | ~ (h10(op1(e11, e15)) = op2(h10(e11), h10(e15))) | ~ (h10(op1(e11, e14)) = op2(h10(e11), h10(e14))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e15)) = op2(h10(e10), h10(e15))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(forward_demodulation, [], [f15076, f5319])).
fof(f15076, plain, (~ (h10(op1(e14, e15)) = op2(h10(e14), e25)) | ~ (h5(e13) = h10(op1(e15, e10))) | ~ (h10(e13) = h10(op1(e15, e11))) | ~ (h10(op1(e15, e12)) = op2(e25, h10(e12))) | ~ (h10(op1(e15, e13)) = op2(e25, h10(e13))) | ~ (h10(e12) = h10(op1(e15, e14))) | ~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e14, e14)) = op2(h10(e14), h10(e14))) | ~ (h10(op1(e14, e13)) = op2(h10(e14), h10(e13))) | ~ (h10(op1(e14, e12)) = op2(h10(e14), h10(e12))) | ~ (h10(op1(e14, e11)) = op2(h10(e14), h10(e11))) | ~ (h10(op1(e14, e10)) = op2(h10(e14), h10(e10))) | ~ (h10(op1(e13, e15)) = op2(h10(e13), h10(e15))) | ~ (h10(op1(e13, e14)) = op2(h10(e13), h10(e14))) | ~ (h10(op1(e13, e13)) = op2(h10(e13), h10(e13))) | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) | ~ (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) | ~ (h10(op1(e12, e15)) = op2(h10(e12), h10(e15))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), h10(e14))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) | ~ (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) | ~ (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) | ~ (h10(op1(e11, e15)) = op2(h10(e11), h10(e15))) | ~ (h10(op1(e11, e14)) = op2(h10(e11), h10(e14))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e15)) = op2(h10(e10), h10(e15))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(forward_demodulation, [], [f15075, f1537])).
fof(f15075, plain, (~ (h5(e13) = h10(op1(e15, e10))) | ~ (h10(e13) = h10(op1(e15, e11))) | ~ (h10(op1(e15, e12)) = op2(e25, h10(e12))) | ~ (h10(op1(e15, e13)) = op2(e25, h10(e13))) | ~ (h10(e12) = h10(op1(e15, e14))) | ~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e14, e15)) = op2(h10(e14), h10(e15))) | ~ (h10(op1(e14, e14)) = op2(h10(e14), h10(e14))) | ~ (h10(op1(e14, e13)) = op2(h10(e14), h10(e13))) | ~ (h10(op1(e14, e12)) = op2(h10(e14), h10(e12))) | ~ (h10(op1(e14, e11)) = op2(h10(e14), h10(e11))) | ~ (h10(op1(e14, e10)) = op2(h10(e14), h10(e10))) | ~ (h10(op1(e13, e15)) = op2(h10(e13), h10(e15))) | ~ (h10(op1(e13, e14)) = op2(h10(e13), h10(e14))) | ~ (h10(op1(e13, e13)) = op2(h10(e13), h10(e13))) | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) | ~ (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) | ~ (h10(op1(e12, e15)) = op2(h10(e12), h10(e15))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), h10(e14))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) | ~ (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) | ~ (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) | ~ (h10(op1(e11, e15)) = op2(h10(e11), h10(e15))) | ~ (h10(op1(e11, e14)) = op2(h10(e11), h10(e14))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e15)) = op2(h10(e10), h10(e15))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(forward_demodulation, [], [f15074, f1510])).
fof(f15074, plain, (~ (op2(e25, e20) = h10(op1(e15, e10))) | ~ (h10(e13) = h10(op1(e15, e11))) | ~ (h10(op1(e15, e12)) = op2(e25, h10(e12))) | ~ (h10(op1(e15, e13)) = op2(e25, h10(e13))) | ~ (h10(e12) = h10(op1(e15, e14))) | ~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e14, e15)) = op2(h10(e14), h10(e15))) | ~ (h10(op1(e14, e14)) = op2(h10(e14), h10(e14))) | ~ (h10(op1(e14, e13)) = op2(h10(e14), h10(e13))) | ~ (h10(op1(e14, e12)) = op2(h10(e14), h10(e12))) | ~ (h10(op1(e14, e11)) = op2(h10(e14), h10(e11))) | ~ (h10(op1(e14, e10)) = op2(h10(e14), h10(e10))) | ~ (h10(op1(e13, e15)) = op2(h10(e13), h10(e15))) | ~ (h10(op1(e13, e14)) = op2(h10(e13), h10(e14))) | ~ (h10(op1(e13, e13)) = op2(h10(e13), h10(e13))) | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) | ~ (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) | ~ (h10(op1(e12, e15)) = op2(h10(e12), h10(e15))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), h10(e14))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) | ~ (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) | ~ (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) | ~ (h10(op1(e11, e15)) = op2(h10(e11), h10(e15))) | ~ (h10(op1(e11, e14)) = op2(h10(e11), h10(e14))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e15)) = op2(h10(e10), h10(e15))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(forward_demodulation, [], [f15073, f1537])).
fof(f15073, plain, (~ (h10(op1(e15, e10)) = op2(h10(e15), e20)) | ~ (h10(e13) = h10(op1(e15, e11))) | ~ (h10(op1(e15, e12)) = op2(e25, h10(e12))) | ~ (h10(op1(e15, e13)) = op2(e25, h10(e13))) | ~ (h10(e12) = h10(op1(e15, e14))) | ~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e14, e15)) = op2(h10(e14), h10(e15))) | ~ (h10(op1(e14, e14)) = op2(h10(e14), h10(e14))) | ~ (h10(op1(e14, e13)) = op2(h10(e14), h10(e13))) | ~ (h10(op1(e14, e12)) = op2(h10(e14), h10(e12))) | ~ (h10(op1(e14, e11)) = op2(h10(e14), h10(e11))) | ~ (h10(op1(e14, e10)) = op2(h10(e14), h10(e10))) | ~ (h10(op1(e13, e15)) = op2(h10(e13), h10(e15))) | ~ (h10(op1(e13, e14)) = op2(h10(e13), h10(e14))) | ~ (h10(op1(e13, e13)) = op2(h10(e13), h10(e13))) | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) | ~ (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) | ~ (h10(op1(e12, e15)) = op2(h10(e12), h10(e15))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), h10(e14))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) | ~ (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) | ~ (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) | ~ (h10(op1(e11, e15)) = op2(h10(e11), h10(e15))) | ~ (h10(op1(e11, e14)) = op2(h10(e11), h10(e14))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e15)) = op2(h10(e10), h10(e15))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(forward_demodulation, [], [f15072, f5365])).
fof(f15072, plain, (~ (h10(e13) = h10(op1(e15, e11))) | ~ (h10(op1(e15, e12)) = op2(e25, h10(e12))) | ~ (h10(op1(e15, e13)) = op2(e25, h10(e13))) | ~ (h10(e12) = h10(op1(e15, e14))) | ~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e15, e10)) = op2(h10(e15), h10(e10))) | ~ (h10(op1(e14, e15)) = op2(h10(e14), h10(e15))) | ~ (h10(op1(e14, e14)) = op2(h10(e14), h10(e14))) | ~ (h10(op1(e14, e13)) = op2(h10(e14), h10(e13))) | ~ (h10(op1(e14, e12)) = op2(h10(e14), h10(e12))) | ~ (h10(op1(e14, e11)) = op2(h10(e14), h10(e11))) | ~ (h10(op1(e14, e10)) = op2(h10(e14), h10(e10))) | ~ (h10(op1(e13, e15)) = op2(h10(e13), h10(e15))) | ~ (h10(op1(e13, e14)) = op2(h10(e13), h10(e14))) | ~ (h10(op1(e13, e13)) = op2(h10(e13), h10(e13))) | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) | ~ (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) | ~ (h10(op1(e12, e15)) = op2(h10(e12), h10(e15))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), h10(e14))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) | ~ (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) | ~ (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) | ~ (h10(op1(e11, e15)) = op2(h10(e11), h10(e15))) | ~ (h10(op1(e11, e14)) = op2(h10(e11), h10(e14))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e15)) = op2(h10(e10), h10(e15))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(forward_demodulation, [], [f15071, f1540])).
fof(f15071, plain, (~ (op2(e25, e21) = h10(op1(e15, e11))) | ~ (h10(op1(e15, e12)) = op2(e25, h10(e12))) | ~ (h10(op1(e15, e13)) = op2(e25, h10(e13))) | ~ (h10(e12) = h10(op1(e15, e14))) | ~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e15, e10)) = op2(h10(e15), h10(e10))) | ~ (h10(op1(e14, e15)) = op2(h10(e14), h10(e15))) | ~ (h10(op1(e14, e14)) = op2(h10(e14), h10(e14))) | ~ (h10(op1(e14, e13)) = op2(h10(e14), h10(e13))) | ~ (h10(op1(e14, e12)) = op2(h10(e14), h10(e12))) | ~ (h10(op1(e14, e11)) = op2(h10(e14), h10(e11))) | ~ (h10(op1(e14, e10)) = op2(h10(e14), h10(e10))) | ~ (h10(op1(e13, e15)) = op2(h10(e13), h10(e15))) | ~ (h10(op1(e13, e14)) = op2(h10(e13), h10(e14))) | ~ (h10(op1(e13, e13)) = op2(h10(e13), h10(e13))) | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) | ~ (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) | ~ (h10(op1(e12, e15)) = op2(h10(e12), h10(e15))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), h10(e14))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) | ~ (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) | ~ (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) | ~ (h10(op1(e11, e15)) = op2(h10(e11), h10(e15))) | ~ (h10(op1(e11, e14)) = op2(h10(e11), h10(e14))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e15)) = op2(h10(e10), h10(e15))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(forward_demodulation, [], [f15070, f1537])).
fof(f15070, plain, (~ (h10(op1(e15, e11)) = op2(h10(e15), e21)) | ~ (h10(op1(e15, e12)) = op2(e25, h10(e12))) | ~ (h10(op1(e15, e13)) = op2(e25, h10(e13))) | ~ (h10(e12) = h10(op1(e15, e14))) | ~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e15, e10)) = op2(h10(e15), h10(e10))) | ~ (h10(op1(e14, e15)) = op2(h10(e14), h10(e15))) | ~ (h10(op1(e14, e14)) = op2(h10(e14), h10(e14))) | ~ (h10(op1(e14, e13)) = op2(h10(e14), h10(e13))) | ~ (h10(op1(e14, e12)) = op2(h10(e14), h10(e12))) | ~ (h10(op1(e14, e11)) = op2(h10(e14), h10(e11))) | ~ (h10(op1(e14, e10)) = op2(h10(e14), h10(e10))) | ~ (h10(op1(e13, e15)) = op2(h10(e13), h10(e15))) | ~ (h10(op1(e13, e14)) = op2(h10(e13), h10(e14))) | ~ (h10(op1(e13, e13)) = op2(h10(e13), h10(e13))) | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) | ~ (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) | ~ (h10(op1(e12, e15)) = op2(h10(e12), h10(e15))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), h10(e14))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) | ~ (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) | ~ (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) | ~ (h10(op1(e11, e15)) = op2(h10(e11), h10(e15))) | ~ (h10(op1(e11, e14)) = op2(h10(e11), h10(e14))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e15)) = op2(h10(e10), h10(e15))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(forward_demodulation, [], [f15069, f1536])).
fof(f15069, plain, (~ (h10(op1(e15, e12)) = op2(e25, h10(e12))) | ~ (h10(op1(e15, e13)) = op2(e25, h10(e13))) | ~ (h10(e12) = h10(op1(e15, e14))) | ~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e15, e11)) = op2(h10(e15), h10(e11))) | ~ (h10(op1(e15, e10)) = op2(h10(e15), h10(e10))) | ~ (h10(op1(e14, e15)) = op2(h10(e14), h10(e15))) | ~ (h10(op1(e14, e14)) = op2(h10(e14), h10(e14))) | ~ (h10(op1(e14, e13)) = op2(h10(e14), h10(e13))) | ~ (h10(op1(e14, e12)) = op2(h10(e14), h10(e12))) | ~ (h10(op1(e14, e11)) = op2(h10(e14), h10(e11))) | ~ (h10(op1(e14, e10)) = op2(h10(e14), h10(e10))) | ~ (h10(op1(e13, e15)) = op2(h10(e13), h10(e15))) | ~ (h10(op1(e13, e14)) = op2(h10(e13), h10(e14))) | ~ (h10(op1(e13, e13)) = op2(h10(e13), h10(e13))) | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) | ~ (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) | ~ (h10(op1(e12, e15)) = op2(h10(e12), h10(e15))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), h10(e14))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) | ~ (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) | ~ (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) | ~ (h10(op1(e11, e15)) = op2(h10(e11), h10(e15))) | ~ (h10(op1(e11, e14)) = op2(h10(e11), h10(e14))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e15)) = op2(h10(e10), h10(e15))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(forward_demodulation, [], [f15068, f1537])).
fof(f15068, plain, (~ (h10(op1(e15, e13)) = op2(e25, h10(e13))) | ~ (h10(e12) = h10(op1(e15, e14))) | ~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e15, e12)) = op2(h10(e15), h10(e12))) | ~ (h10(op1(e15, e11)) = op2(h10(e15), h10(e11))) | ~ (h10(op1(e15, e10)) = op2(h10(e15), h10(e10))) | ~ (h10(op1(e14, e15)) = op2(h10(e14), h10(e15))) | ~ (h10(op1(e14, e14)) = op2(h10(e14), h10(e14))) | ~ (h10(op1(e14, e13)) = op2(h10(e14), h10(e13))) | ~ (h10(op1(e14, e12)) = op2(h10(e14), h10(e12))) | ~ (h10(op1(e14, e11)) = op2(h10(e14), h10(e11))) | ~ (h10(op1(e14, e10)) = op2(h10(e14), h10(e10))) | ~ (h10(op1(e13, e15)) = op2(h10(e13), h10(e15))) | ~ (h10(op1(e13, e14)) = op2(h10(e13), h10(e14))) | ~ (h10(op1(e13, e13)) = op2(h10(e13), h10(e13))) | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) | ~ (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) | ~ (h10(op1(e12, e15)) = op2(h10(e12), h10(e15))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), h10(e14))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) | ~ (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) | ~ (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) | ~ (h10(op1(e11, e15)) = op2(h10(e11), h10(e15))) | ~ (h10(op1(e11, e14)) = op2(h10(e11), h10(e14))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e15)) = op2(h10(e10), h10(e15))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(forward_demodulation, [], [f15067, f1537])).
fof(f15067, plain, (~ (h10(e12) = h10(op1(e15, e14))) | ~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e15, e13)) = op2(h10(e15), h10(e13))) | ~ (h10(op1(e15, e12)) = op2(h10(e15), h10(e12))) | ~ (h10(op1(e15, e11)) = op2(h10(e15), h10(e11))) | ~ (h10(op1(e15, e10)) = op2(h10(e15), h10(e10))) | ~ (h10(op1(e14, e15)) = op2(h10(e14), h10(e15))) | ~ (h10(op1(e14, e14)) = op2(h10(e14), h10(e14))) | ~ (h10(op1(e14, e13)) = op2(h10(e14), h10(e13))) | ~ (h10(op1(e14, e12)) = op2(h10(e14), h10(e12))) | ~ (h10(op1(e14, e11)) = op2(h10(e14), h10(e11))) | ~ (h10(op1(e14, e10)) = op2(h10(e14), h10(e10))) | ~ (h10(op1(e13, e15)) = op2(h10(e13), h10(e15))) | ~ (h10(op1(e13, e14)) = op2(h10(e13), h10(e14))) | ~ (h10(op1(e13, e13)) = op2(h10(e13), h10(e13))) | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) | ~ (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) | ~ (h10(op1(e12, e15)) = op2(h10(e12), h10(e15))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), h10(e14))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) | ~ (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) | ~ (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) | ~ (h10(op1(e11, e15)) = op2(h10(e11), h10(e15))) | ~ (h10(op1(e11, e14)) = op2(h10(e11), h10(e14))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e15)) = op2(h10(e10), h10(e15))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(forward_demodulation, [], [f15066, f5339])).
fof(f15066, plain, (~ (op2(e25, e24) = h10(op1(e15, e14))) | ~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e15, e13)) = op2(h10(e15), h10(e13))) | ~ (h10(op1(e15, e12)) = op2(h10(e15), h10(e12))) | ~ (h10(op1(e15, e11)) = op2(h10(e15), h10(e11))) | ~ (h10(op1(e15, e10)) = op2(h10(e15), h10(e10))) | ~ (h10(op1(e14, e15)) = op2(h10(e14), h10(e15))) | ~ (h10(op1(e14, e14)) = op2(h10(e14), h10(e14))) | ~ (h10(op1(e14, e13)) = op2(h10(e14), h10(e13))) | ~ (h10(op1(e14, e12)) = op2(h10(e14), h10(e12))) | ~ (h10(op1(e14, e11)) = op2(h10(e14), h10(e11))) | ~ (h10(op1(e14, e10)) = op2(h10(e14), h10(e10))) | ~ (h10(op1(e13, e15)) = op2(h10(e13), h10(e15))) | ~ (h10(op1(e13, e14)) = op2(h10(e13), h10(e14))) | ~ (h10(op1(e13, e13)) = op2(h10(e13), h10(e13))) | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) | ~ (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) | ~ (h10(op1(e12, e15)) = op2(h10(e12), h10(e15))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), h10(e14))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) | ~ (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) | ~ (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) | ~ (h10(op1(e11, e15)) = op2(h10(e11), h10(e15))) | ~ (h10(op1(e11, e14)) = op2(h10(e11), h10(e14))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e15)) = op2(h10(e10), h10(e15))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(forward_demodulation, [], [f15065, f1537])).
fof(f15065, plain, (~ (h10(op1(e15, e14)) = op2(h10(e15), e24)) | ~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e15, e13)) = op2(h10(e15), h10(e13))) | ~ (h10(op1(e15, e12)) = op2(h10(e15), h10(e12))) | ~ (h10(op1(e15, e11)) = op2(h10(e15), h10(e11))) | ~ (h10(op1(e15, e10)) = op2(h10(e15), h10(e10))) | ~ (h10(op1(e14, e15)) = op2(h10(e14), h10(e15))) | ~ (h10(op1(e14, e14)) = op2(h10(e14), h10(e14))) | ~ (h10(op1(e14, e13)) = op2(h10(e14), h10(e13))) | ~ (h10(op1(e14, e12)) = op2(h10(e14), h10(e12))) | ~ (h10(op1(e14, e11)) = op2(h10(e14), h10(e11))) | ~ (h10(op1(e14, e10)) = op2(h10(e14), h10(e10))) | ~ (h10(op1(e13, e15)) = op2(h10(e13), h10(e15))) | ~ (h10(op1(e13, e14)) = op2(h10(e13), h10(e14))) | ~ (h10(op1(e13, e13)) = op2(h10(e13), h10(e13))) | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) | ~ (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) | ~ (h10(op1(e12, e15)) = op2(h10(e12), h10(e15))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), h10(e14))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) | ~ (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) | ~ (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) | ~ (h10(op1(e11, e15)) = op2(h10(e11), h10(e15))) | ~ (h10(op1(e11, e14)) = op2(h10(e11), h10(e14))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e15)) = op2(h10(e10), h10(e15))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(forward_demodulation, [], [f15064, f5319])).
fof(f15064, plain, (~ (op2(e25, e25) = h10(op1(e15, e15))) | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e15, e14)) = op2(h10(e15), h10(e14))) | ~ (h10(op1(e15, e13)) = op2(h10(e15), h10(e13))) | ~ (h10(op1(e15, e12)) = op2(h10(e15), h10(e12))) | ~ (h10(op1(e15, e11)) = op2(h10(e15), h10(e11))) | ~ (h10(op1(e15, e10)) = op2(h10(e15), h10(e10))) | ~ (h10(op1(e14, e15)) = op2(h10(e14), h10(e15))) | ~ (h10(op1(e14, e14)) = op2(h10(e14), h10(e14))) | ~ (h10(op1(e14, e13)) = op2(h10(e14), h10(e13))) | ~ (h10(op1(e14, e12)) = op2(h10(e14), h10(e12))) | ~ (h10(op1(e14, e11)) = op2(h10(e14), h10(e11))) | ~ (h10(op1(e14, e10)) = op2(h10(e14), h10(e10))) | ~ (h10(op1(e13, e15)) = op2(h10(e13), h10(e15))) | ~ (h10(op1(e13, e14)) = op2(h10(e13), h10(e14))) | ~ (h10(op1(e13, e13)) = op2(h10(e13), h10(e13))) | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) | ~ (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) | ~ (h10(op1(e12, e15)) = op2(h10(e12), h10(e15))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), h10(e14))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) | ~ (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) | ~ (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) | ~ (h10(op1(e11, e15)) = op2(h10(e11), h10(e15))) | ~ (h10(op1(e11, e14)) = op2(h10(e11), h10(e14))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e15)) = op2(h10(e10), h10(e15))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(forward_demodulation, [], [f15063, f1537])).
fof(f15063, plain, (sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e15, e15)) = op2(h10(e15), h10(e15))) | ~ (h10(op1(e15, e14)) = op2(h10(e15), h10(e14))) | ~ (h10(op1(e15, e13)) = op2(h10(e15), h10(e13))) | ~ (h10(op1(e15, e12)) = op2(h10(e15), h10(e12))) | ~ (h10(op1(e15, e11)) = op2(h10(e15), h10(e11))) | ~ (h10(op1(e15, e10)) = op2(h10(e15), h10(e10))) | ~ (h10(op1(e14, e15)) = op2(h10(e14), h10(e15))) | ~ (h10(op1(e14, e14)) = op2(h10(e14), h10(e14))) | ~ (h10(op1(e14, e13)) = op2(h10(e14), h10(e13))) | ~ (h10(op1(e14, e12)) = op2(h10(e14), h10(e12))) | ~ (h10(op1(e14, e11)) = op2(h10(e14), h10(e11))) | ~ (h10(op1(e14, e10)) = op2(h10(e14), h10(e10))) | ~ (h10(op1(e13, e15)) = op2(h10(e13), h10(e15))) | ~ (h10(op1(e13, e14)) = op2(h10(e13), h10(e14))) | ~ (h10(op1(e13, e13)) = op2(h10(e13), h10(e13))) | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) | ~ (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) | ~ (h10(op1(e12, e15)) = op2(h10(e12), h10(e15))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), h10(e14))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) | ~ (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) | ~ (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) | ~ (h10(op1(e11, e15)) = op2(h10(e11), h10(e15))) | ~ (h10(op1(e11, e14)) = op2(h10(e11), h10(e14))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e15)) = op2(h10(e10), h10(e15))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(subsumption_resolution, [], [f15062, f7977])).
fof(f7977, plain, ~ sP49, inference(subsumption_resolution, [], [f2266, f5319])).
fof(f2266, plain, (~ (e24 = h10(e14)) | ~ sP49), inference(cnf_transformation, [], [f314])).
fof(f314, plain, ((~ (e24 = h10(e15)) & ~ (e24 = h10(e14)) & ~ (e24 = h10(e13)) & ~ (e24 = h10(e12)) & ~ (e24 = h10(e11)) & ~ (e24 = h10(e10))) | ~ sP49), inference(nnf_transformation, [], [f112])).
fof(f112, plain, ((~ (e24 = h10(e15)) & ~ (e24 = h10(e14)) & ~ (e24 = h10(e13)) & ~ (e24 = h10(e12)) & ~ (e24 = h10(e11)) & ~ (e24 = h10(e10))) | ~ sP49), inference(usedef, [], [e112])).
fof(e112, plain, (sP49 <=> (~ (e24 = h10(e15)) & ~ (e24 = h10(e14)) & ~ (e24 = h10(e13)) & ~ (e24 = h10(e12)) & ~ (e24 = h10(e11)) & ~ (e24 = h10(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP49])])).
fof(f15062, plain, (sP49 | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e15, e15)) = op2(h10(e15), h10(e15))) | ~ (h10(op1(e15, e14)) = op2(h10(e15), h10(e14))) | ~ (h10(op1(e15, e13)) = op2(h10(e15), h10(e13))) | ~ (h10(op1(e15, e12)) = op2(h10(e15), h10(e12))) | ~ (h10(op1(e15, e11)) = op2(h10(e15), h10(e11))) | ~ (h10(op1(e15, e10)) = op2(h10(e15), h10(e10))) | ~ (h10(op1(e14, e15)) = op2(h10(e14), h10(e15))) | ~ (h10(op1(e14, e14)) = op2(h10(e14), h10(e14))) | ~ (h10(op1(e14, e13)) = op2(h10(e14), h10(e13))) | ~ (h10(op1(e14, e12)) = op2(h10(e14), h10(e12))) | ~ (h10(op1(e14, e11)) = op2(h10(e14), h10(e11))) | ~ (h10(op1(e14, e10)) = op2(h10(e14), h10(e10))) | ~ (h10(op1(e13, e15)) = op2(h10(e13), h10(e15))) | ~ (h10(op1(e13, e14)) = op2(h10(e13), h10(e14))) | ~ (h10(op1(e13, e13)) = op2(h10(e13), h10(e13))) | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) | ~ (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) | ~ (h10(op1(e12, e15)) = op2(h10(e12), h10(e15))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), h10(e14))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) | ~ (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) | ~ (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) | ~ (h10(op1(e11, e15)) = op2(h10(e11), h10(e15))) | ~ (h10(op1(e11, e14)) = op2(h10(e11), h10(e14))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e15)) = op2(h10(e10), h10(e15))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(subsumption_resolution, [], [f2621, f1537])).
fof(f2621, plain, (~ (e25 = h10(e15)) | sP49 | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e15, e15)) = op2(h10(e15), h10(e15))) | ~ (h10(op1(e15, e14)) = op2(h10(e15), h10(e14))) | ~ (h10(op1(e15, e13)) = op2(h10(e15), h10(e13))) | ~ (h10(op1(e15, e12)) = op2(h10(e15), h10(e12))) | ~ (h10(op1(e15, e11)) = op2(h10(e15), h10(e11))) | ~ (h10(op1(e15, e10)) = op2(h10(e15), h10(e10))) | ~ (h10(op1(e14, e15)) = op2(h10(e14), h10(e15))) | ~ (h10(op1(e14, e14)) = op2(h10(e14), h10(e14))) | ~ (h10(op1(e14, e13)) = op2(h10(e14), h10(e13))) | ~ (h10(op1(e14, e12)) = op2(h10(e14), h10(e12))) | ~ (h10(op1(e14, e11)) = op2(h10(e14), h10(e11))) | ~ (h10(op1(e14, e10)) = op2(h10(e14), h10(e10))) | ~ (h10(op1(e13, e15)) = op2(h10(e13), h10(e15))) | ~ (h10(op1(e13, e14)) = op2(h10(e13), h10(e14))) | ~ (h10(op1(e13, e13)) = op2(h10(e13), h10(e13))) | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) | ~ (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) | ~ (h10(op1(e12, e15)) = op2(h10(e12), h10(e15))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), h10(e14))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) | ~ (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) | ~ (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) | ~ (h10(op1(e11, e15)) = op2(h10(e11), h10(e15))) | ~ (h10(op1(e11, e14)) = op2(h10(e11), h10(e14))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e15)) = op2(h10(e10), h10(e15))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))), inference(cnf_transformation, [], [f213])).
fof(f213, plain, (((~ (e25 = h30(e15)) & ~ (e25 = h30(e14)) & ~ (e25 = h30(e13)) & ~ (e25 = h30(e12)) & ~ (e25 = h30(e11)) & ~ (e25 = h30(e10))) | sP149 | sP148 | sP147 | sP146 | sP145 | ~ (h30(op1(e15, e15)) = op2(h30(e15), h30(e15))) | ~ (h30(op1(e15, e14)) = op2(h30(e15), h30(e14))) | ~ (h30(op1(e15, e13)) = op2(h30(e15), h30(e13))) | ~ (h30(op1(e15, e12)) = op2(h30(e15), h30(e12))) | ~ (h30(op1(e15, e11)) = op2(h30(e15), h30(e11))) | ~ (h30(op1(e15, e10)) = op2(h30(e15), h30(e10))) | ~ (h30(op1(e14, e15)) = op2(h30(e14), h30(e15))) | ~ (h30(op1(e14, e14)) = op2(h30(e14), h30(e14))) | ~ (h30(op1(e14, e13)) = op2(h30(e14), h30(e13))) | ~ (h30(op1(e14, e12)) = op2(h30(e14), h30(e12))) | ~ (h30(op1(e14, e11)) = op2(h30(e14), h30(e11))) | ~ (h30(op1(e14, e10)) = op2(h30(e14), h30(e10))) | ~ (h30(op1(e13, e15)) = op2(h30(e13), h30(e15))) | ~ (h30(op1(e13, e14)) = op2(h30(e13), h30(e14))) | ~ (h30(op1(e13, e13)) = op2(h30(e13), h30(e13))) | ~ (h30(op1(e13, e12)) = op2(h30(e13), h30(e12))) | ~ (h30(op1(e13, e11)) = op2(h30(e13), h30(e11))) | ~ (h30(op1(e13, e10)) = op2(h30(e13), h30(e10))) | ~ (h30(op1(e12, e15)) = op2(h30(e12), h30(e15))) | ~ (h30(op1(e12, e14)) = op2(h30(e12), h30(e14))) | ~ (h30(op1(e12, e13)) = op2(h30(e12), h30(e13))) | ~ (h30(op1(e12, e12)) = op2(h30(e12), h30(e12))) | ~ (h30(op1(e12, e11)) = op2(h30(e12), h30(e11))) | ~ (h30(op1(e12, e10)) = op2(h30(e12), h30(e10))) | ~ (h30(op1(e11, e15)) = op2(h30(e11), h30(e15))) | ~ (h30(op1(e11, e14)) = op2(h30(e11), h30(e14))) | ~ (h30(op1(e11, e13)) = op2(h30(e11), h30(e13))) | ~ (h30(op1(e11, e12)) = op2(h30(e11), h30(e12))) | ~ (h30(op1(e11, e11)) = op2(h30(e11), h30(e11))) | ~ (h30(op1(e11, e10)) = op2(h30(e11), h30(e10))) | ~ (h30(op1(e10, e15)) = op2(h30(e10), h30(e15))) | ~ (h30(op1(e10, e14)) = op2(h30(e10), h30(e14))) | ~ (h30(op1(e10, e13)) = op2(h30(e10), h30(e13))) | ~ (h30(op1(e10, e12)) = op2(h30(e10), h30(e12))) | ~ (h30(op1(e10, e11)) = op2(h30(e10), h30(e11))) | ~ (h30(op1(e10, e10)) = op2(h30(e10), h30(e10)))) & ((~ (e25 = h29(e15)) & ~ (e25 = h29(e14)) & ~ (e25 = h29(e13)) & ~ (e25 = h29(e12)) & ~ (e25 = h29(e11)) & ~ (e25 = h29(e10))) | sP144 | sP143 | sP142 | sP141 | sP140 | ~ (h29(op1(e15, e15)) = op2(h29(e15), h29(e15))) | ~ (h29(op1(e15, e14)) = op2(h29(e15), h29(e14))) | ~ (h29(op1(e15, e13)) = op2(h29(e15), h29(e13))) | ~ (h29(op1(e15, e12)) = op2(h29(e15), h29(e12))) | ~ (h29(op1(e15, e11)) = op2(h29(e15), h29(e11))) | ~ (h29(op1(e15, e10)) = op2(h29(e15), h29(e10))) | ~ (h29(op1(e14, e15)) = op2(h29(e14), h29(e15))) | ~ (h29(op1(e14, e14)) = op2(h29(e14), h29(e14))) | ~ (h29(op1(e14, e13)) = op2(h29(e14), h29(e13))) | ~ (h29(op1(e14, e12)) = op2(h29(e14), h29(e12))) | ~ (h29(op1(e14, e11)) = op2(h29(e14), h29(e11))) | ~ (h29(op1(e14, e10)) = op2(h29(e14), h29(e10))) | ~ (h29(op1(e13, e15)) = op2(h29(e13), h29(e15))) | ~ (h29(op1(e13, e14)) = op2(h29(e13), h29(e14))) | ~ (h29(op1(e13, e13)) = op2(h29(e13), h29(e13))) | ~ (h29(op1(e13, e12)) = op2(h29(e13), h29(e12))) | ~ (h29(op1(e13, e11)) = op2(h29(e13), h29(e11))) | ~ (h29(op1(e13, e10)) = op2(h29(e13), h29(e10))) | ~ (h29(op1(e12, e15)) = op2(h29(e12), h29(e15))) | ~ (h29(op1(e12, e14)) = op2(h29(e12), h29(e14))) | ~ (h29(op1(e12, e13)) = op2(h29(e12), h29(e13))) | ~ (h29(op1(e12, e12)) = op2(h29(e12), h29(e12))) | ~ (h29(op1(e12, e11)) = op2(h29(e12), h29(e11))) | ~ (h29(op1(e12, e10)) = op2(h29(e12), h29(e10))) | ~ (h29(op1(e11, e15)) = op2(h29(e11), h29(e15))) | ~ (h29(op1(e11, e14)) = op2(h29(e11), h29(e14))) | ~ (h29(op1(e11, e13)) = op2(h29(e11), h29(e13))) | ~ (h29(op1(e11, e12)) = op2(h29(e11), h29(e12))) | ~ (h29(op1(e11, e11)) = op2(h29(e11), h29(e11))) | ~ (h29(op1(e11, e10)) = op2(h29(e11), h29(e10))) | ~ (h29(op1(e10, e15)) = op2(h29(e10), h29(e15))) | ~ (h29(op1(e10, e14)) = op2(h29(e10), h29(e14))) | ~ (h29(op1(e10, e13)) = op2(h29(e10), h29(e13))) | ~ (h29(op1(e10, e12)) = op2(h29(e10), h29(e12))) | ~ (h29(op1(e10, e11)) = op2(h29(e10), h29(e11))) | ~ (h29(op1(e10, e10)) = op2(h29(e10), h29(e10)))) & ((~ (e25 = h28(e15)) & ~ (e25 = h28(e14)) & ~ (e25 = h28(e13)) & ~ (e25 = h28(e12)) & ~ (e25 = h28(e11)) & ~ (e25 = h28(e10))) | sP139 | sP138 | sP137 | sP136 | sP135 | ~ (h28(op1(e15, e15)) = op2(h28(e15), h28(e15))) | ~ (h28(op1(e15, e14)) = op2(h28(e15), h28(e14))) | ~ (h28(op1(e15, e13)) = op2(h28(e15), h28(e13))) | ~ (h28(op1(e15, e12)) = op2(h28(e15), h28(e12))) | ~ (h28(op1(e15, e11)) = op2(h28(e15), h28(e11))) | ~ (h28(op1(e15, e10)) = op2(h28(e15), h28(e10))) | ~ (h28(op1(e14, e15)) = op2(h28(e14), h28(e15))) | ~ (h28(op1(e14, e14)) = op2(h28(e14), h28(e14))) | ~ (h28(op1(e14, e13)) = op2(h28(e14), h28(e13))) | ~ (h28(op1(e14, e12)) = op2(h28(e14), h28(e12))) | ~ (h28(op1(e14, e11)) = op2(h28(e14), h28(e11))) | ~ (h28(op1(e14, e10)) = op2(h28(e14), h28(e10))) | ~ (h28(op1(e13, e15)) = op2(h28(e13), h28(e15))) | ~ (h28(op1(e13, e14)) = op2(h28(e13), h28(e14))) | ~ (h28(op1(e13, e13)) = op2(h28(e13), h28(e13))) | ~ (h28(op1(e13, e12)) = op2(h28(e13), h28(e12))) | ~ (h28(op1(e13, e11)) = op2(h28(e13), h28(e11))) | ~ (h28(op1(e13, e10)) = op2(h28(e13), h28(e10))) | ~ (h28(op1(e12, e15)) = op2(h28(e12), h28(e15))) | ~ (h28(op1(e12, e14)) = op2(h28(e12), h28(e14))) | ~ (h28(op1(e12, e13)) = op2(h28(e12), h28(e13))) | ~ (h28(op1(e12, e12)) = op2(h28(e12), h28(e12))) | ~ (h28(op1(e12, e11)) = op2(h28(e12), h28(e11))) | ~ (h28(op1(e12, e10)) = op2(h28(e12), h28(e10))) | ~ (h28(op1(e11, e15)) = op2(h28(e11), h28(e15))) | ~ (h28(op1(e11, e14)) = op2(h28(e11), h28(e14))) | ~ (h28(op1(e11, e13)) = op2(h28(e11), h28(e13))) | ~ (h28(op1(e11, e12)) = op2(h28(e11), h28(e12))) | ~ (h28(op1(e11, e11)) = op2(h28(e11), h28(e11))) | ~ (h28(op1(e11, e10)) = op2(h28(e11), h28(e10))) | ~ (h28(op1(e10, e15)) = op2(h28(e10), h28(e15))) | ~ (h28(op1(e10, e14)) = op2(h28(e10), h28(e14))) | ~ (h28(op1(e10, e13)) = op2(h28(e10), h28(e13))) | ~ (h28(op1(e10, e12)) = op2(h28(e10), h28(e12))) | ~ (h28(op1(e10, e11)) = op2(h28(e10), h28(e11))) | ~ (h28(op1(e10, e10)) = op2(h28(e10), h28(e10)))) & ((~ (e25 = h27(e15)) & ~ (e25 = h27(e14)) & ~ (e25 = h27(e13)) & ~ (e25 = h27(e12)) & ~ (e25 = h27(e11)) & ~ (e25 = h27(e10))) | sP134 | sP133 | sP132 | sP131 | sP130 | ~ (h27(op1(e15, e15)) = op2(h27(e15), h27(e15))) | ~ (h27(op1(e15, e14)) = op2(h27(e15), h27(e14))) | ~ (h27(op1(e15, e13)) = op2(h27(e15), h27(e13))) | ~ (h27(op1(e15, e12)) = op2(h27(e15), h27(e12))) | ~ (h27(op1(e15, e11)) = op2(h27(e15), h27(e11))) | ~ (h27(op1(e15, e10)) = op2(h27(e15), h27(e10))) | ~ (h27(op1(e14, e15)) = op2(h27(e14), h27(e15))) | ~ (h27(op1(e14, e14)) = op2(h27(e14), h27(e14))) | ~ (h27(op1(e14, e13)) = op2(h27(e14), h27(e13))) | ~ (h27(op1(e14, e12)) = op2(h27(e14), h27(e12))) | ~ (h27(op1(e14, e11)) = op2(h27(e14), h27(e11))) | ~ (h27(op1(e14, e10)) = op2(h27(e14), h27(e10))) | ~ (h27(op1(e13, e15)) = op2(h27(e13), h27(e15))) | ~ (h27(op1(e13, e14)) = op2(h27(e13), h27(e14))) | ~ (h27(op1(e13, e13)) = op2(h27(e13), h27(e13))) | ~ (h27(op1(e13, e12)) = op2(h27(e13), h27(e12))) | ~ (h27(op1(e13, e11)) = op2(h27(e13), h27(e11))) | ~ (h27(op1(e13, e10)) = op2(h27(e13), h27(e10))) | ~ (h27(op1(e12, e15)) = op2(h27(e12), h27(e15))) | ~ (h27(op1(e12, e14)) = op2(h27(e12), h27(e14))) | ~ (h27(op1(e12, e13)) = op2(h27(e12), h27(e13))) | ~ (h27(op1(e12, e12)) = op2(h27(e12), h27(e12))) | ~ (h27(op1(e12, e11)) = op2(h27(e12), h27(e11))) | ~ (h27(op1(e12, e10)) = op2(h27(e12), h27(e10))) | ~ (h27(op1(e11, e15)) = op2(h27(e11), h27(e15))) | ~ (h27(op1(e11, e14)) = op2(h27(e11), h27(e14))) | ~ (h27(op1(e11, e13)) = op2(h27(e11), h27(e13))) | ~ (h27(op1(e11, e12)) = op2(h27(e11), h27(e12))) | ~ (h27(op1(e11, e11)) = op2(h27(e11), h27(e11))) | ~ (h27(op1(e11, e10)) = op2(h27(e11), h27(e10))) | ~ (h27(op1(e10, e15)) = op2(h27(e10), h27(e15))) | ~ (h27(op1(e10, e14)) = op2(h27(e10), h27(e14))) | ~ (h27(op1(e10, e13)) = op2(h27(e10), h27(e13))) | ~ (h27(op1(e10, e12)) = op2(h27(e10), h27(e12))) | ~ (h27(op1(e10, e11)) = op2(h27(e10), h27(e11))) | ~ (h27(op1(e10, e10)) = op2(h27(e10), h27(e10)))) & ((~ (e25 = h26(e15)) & ~ (e25 = h26(e14)) & ~ (e25 = h26(e13)) & ~ (e25 = h26(e12)) & ~ (e25 = h26(e11)) & ~ (e25 = h26(e10))) | sP129 | sP128 | sP127 | sP126 | sP125 | ~ (h26(op1(e15, e15)) = op2(h26(e15), h26(e15))) | ~ (h26(op1(e15, e14)) = op2(h26(e15), h26(e14))) | ~ (h26(op1(e15, e13)) = op2(h26(e15), h26(e13))) | ~ (h26(op1(e15, e12)) = op2(h26(e15), h26(e12))) | ~ (h26(op1(e15, e11)) = op2(h26(e15), h26(e11))) | ~ (h26(op1(e15, e10)) = op2(h26(e15), h26(e10))) | ~ (h26(op1(e14, e15)) = op2(h26(e14), h26(e15))) | ~ (h26(op1(e14, e14)) = op2(h26(e14), h26(e14))) | ~ (h26(op1(e14, e13)) = op2(h26(e14), h26(e13))) | ~ (h26(op1(e14, e12)) = op2(h26(e14), h26(e12))) | ~ (h26(op1(e14, e11)) = op2(h26(e14), h26(e11))) | ~ (h26(op1(e14, e10)) = op2(h26(e14), h26(e10))) | ~ (h26(op1(e13, e15)) = op2(h26(e13), h26(e15))) | ~ (h26(op1(e13, e14)) = op2(h26(e13), h26(e14))) | ~ (h26(op1(e13, e13)) = op2(h26(e13), h26(e13))) | ~ (h26(op1(e13, e12)) = op2(h26(e13), h26(e12))) | ~ (h26(op1(e13, e11)) = op2(h26(e13), h26(e11))) | ~ (h26(op1(e13, e10)) = op2(h26(e13), h26(e10))) | ~ (h26(op1(e12, e15)) = op2(h26(e12), h26(e15))) | ~ (h26(op1(e12, e14)) = op2(h26(e12), h26(e14))) | ~ (h26(op1(e12, e13)) = op2(h26(e12), h26(e13))) | ~ (h26(op1(e12, e12)) = op2(h26(e12), h26(e12))) | ~ (h26(op1(e12, e11)) = op2(h26(e12), h26(e11))) | ~ (h26(op1(e12, e10)) = op2(h26(e12), h26(e10))) | ~ (h26(op1(e11, e15)) = op2(h26(e11), h26(e15))) | ~ (h26(op1(e11, e14)) = op2(h26(e11), h26(e14))) | ~ (h26(op1(e11, e13)) = op2(h26(e11), h26(e13))) | ~ (h26(op1(e11, e12)) = op2(h26(e11), h26(e12))) | ~ (h26(op1(e11, e11)) = op2(h26(e11), h26(e11))) | ~ (h26(op1(e11, e10)) = op2(h26(e11), h26(e10))) | ~ (h26(op1(e10, e15)) = op2(h26(e10), h26(e15))) | ~ (h26(op1(e10, e14)) = op2(h26(e10), h26(e14))) | ~ (h26(op1(e10, e13)) = op2(h26(e10), h26(e13))) | ~ (h26(op1(e10, e12)) = op2(h26(e10), h26(e12))) | ~ (h26(op1(e10, e11)) = op2(h26(e10), h26(e11))) | ~ (h26(op1(e10, e10)) = op2(h26(e10), h26(e10)))) & ((~ (e25 = h25(e15)) & ~ (e25 = h25(e14)) & ~ (e25 = h25(e13)) & ~ (e25 = h25(e12)) & ~ (e25 = h25(e11)) & ~ (e25 = h25(e10))) | sP124 | sP123 | sP122 | sP121 | sP120 | ~ (h25(op1(e15, e15)) = op2(h25(e15), h25(e15))) | ~ (h25(op1(e15, e14)) = op2(h25(e15), h25(e14))) | ~ (h25(op1(e15, e13)) = op2(h25(e15), h25(e13))) | ~ (h25(op1(e15, e12)) = op2(h25(e15), h25(e12))) | ~ (h25(op1(e15, e11)) = op2(h25(e15), h25(e11))) | ~ (h25(op1(e15, e10)) = op2(h25(e15), h25(e10))) | ~ (h25(op1(e14, e15)) = op2(h25(e14), h25(e15))) | ~ (h25(op1(e14, e14)) = op2(h25(e14), h25(e14))) | ~ (h25(op1(e14, e13)) = op2(h25(e14), h25(e13))) | ~ (h25(op1(e14, e12)) = op2(h25(e14), h25(e12))) | ~ (h25(op1(e14, e11)) = op2(h25(e14), h25(e11))) | ~ (h25(op1(e14, e10)) = op2(h25(e14), h25(e10))) | ~ (h25(op1(e13, e15)) = op2(h25(e13), h25(e15))) | ~ (h25(op1(e13, e14)) = op2(h25(e13), h25(e14))) | ~ (h25(op1(e13, e13)) = op2(h25(e13), h25(e13))) | ~ (h25(op1(e13, e12)) = op2(h25(e13), h25(e12))) | ~ (h25(op1(e13, e11)) = op2(h25(e13), h25(e11))) | ~ (h25(op1(e13, e10)) = op2(h25(e13), h25(e10))) | ~ (h25(op1(e12, e15)) = op2(h25(e12), h25(e15))) | ~ (h25(op1(e12, e14)) = op2(h25(e12), h25(e14))) | ~ (h25(op1(e12, e13)) = op2(h25(e12), h25(e13))) | ~ (h25(op1(e12, e12)) = op2(h25(e12), h25(e12))) | ~ (h25(op1(e12, e11)) = op2(h25(e12), h25(e11))) | ~ (h25(op1(e12, e10)) = op2(h25(e12), h25(e10))) | ~ (h25(op1(e11, e15)) = op2(h25(e11), h25(e15))) | ~ (h25(op1(e11, e14)) = op2(h25(e11), h25(e14))) | ~ (h25(op1(e11, e13)) = op2(h25(e11), h25(e13))) | ~ (h25(op1(e11, e12)) = op2(h25(e11), h25(e12))) | ~ (h25(op1(e11, e11)) = op2(h25(e11), h25(e11))) | ~ (h25(op1(e11, e10)) = op2(h25(e11), h25(e10))) | ~ (h25(op1(e10, e15)) = op2(h25(e10), h25(e15))) | ~ (h25(op1(e10, e14)) = op2(h25(e10), h25(e14))) | ~ (h25(op1(e10, e13)) = op2(h25(e10), h25(e13))) | ~ (h25(op1(e10, e12)) = op2(h25(e10), h25(e12))) | ~ (h25(op1(e10, e11)) = op2(h25(e10), h25(e11))) | ~ (h25(op1(e10, e10)) = op2(h25(e10), h25(e10)))) & ((~ (e25 = h24(e15)) & ~ (e25 = h24(e14)) & ~ (e25 = h24(e13)) & ~ (e25 = h24(e12)) & ~ (e25 = h24(e11)) & ~ (e25 = h24(e10))) | sP119 | sP118 | sP117 | sP116 | sP115 | ~ (h24(op1(e15, e15)) = op2(h24(e15), h24(e15))) | ~ (h24(op1(e15, e14)) = op2(h24(e15), h24(e14))) | ~ (h24(op1(e15, e13)) = op2(h24(e15), h24(e13))) | ~ (h24(op1(e15, e12)) = op2(h24(e15), h24(e12))) | ~ (h24(op1(e15, e11)) = op2(h24(e15), h24(e11))) | ~ (h24(op1(e15, e10)) = op2(h24(e15), h24(e10))) | ~ (h24(op1(e14, e15)) = op2(h24(e14), h24(e15))) | ~ (h24(op1(e14, e14)) = op2(h24(e14), h24(e14))) | ~ (h24(op1(e14, e13)) = op2(h24(e14), h24(e13))) | ~ (h24(op1(e14, e12)) = op2(h24(e14), h24(e12))) | ~ (h24(op1(e14, e11)) = op2(h24(e14), h24(e11))) | ~ (h24(op1(e14, e10)) = op2(h24(e14), h24(e10))) | ~ (h24(op1(e13, e15)) = op2(h24(e13), h24(e15))) | ~ (h24(op1(e13, e14)) = op2(h24(e13), h24(e14))) | ~ (h24(op1(e13, e13)) = op2(h24(e13), h24(e13))) | ~ (h24(op1(e13, e12)) = op2(h24(e13), h24(e12))) | ~ (h24(op1(e13, e11)) = op2(h24(e13), h24(e11))) | ~ (h24(op1(e13, e10)) = op2(h24(e13), h24(e10))) | ~ (h24(op1(e12, e15)) = op2(h24(e12), h24(e15))) | ~ (h24(op1(e12, e14)) = op2(h24(e12), h24(e14))) | ~ (h24(op1(e12, e13)) = op2(h24(e12), h24(e13))) | ~ (h24(op1(e12, e12)) = op2(h24(e12), h24(e12))) | ~ (h24(op1(e12, e11)) = op2(h24(e12), h24(e11))) | ~ (h24(op1(e12, e10)) = op2(h24(e12), h24(e10))) | ~ (h24(op1(e11, e15)) = op2(h24(e11), h24(e15))) | ~ (h24(op1(e11, e14)) = op2(h24(e11), h24(e14))) | ~ (h24(op1(e11, e13)) = op2(h24(e11), h24(e13))) | ~ (h24(op1(e11, e12)) = op2(h24(e11), h24(e12))) | ~ (h24(op1(e11, e11)) = op2(h24(e11), h24(e11))) | ~ (h24(op1(e11, e10)) = op2(h24(e11), h24(e10))) | ~ (h24(op1(e10, e15)) = op2(h24(e10), h24(e15))) | ~ (h24(op1(e10, e14)) = op2(h24(e10), h24(e14))) | ~ (h24(op1(e10, e13)) = op2(h24(e10), h24(e13))) | ~ (h24(op1(e10, e12)) = op2(h24(e10), h24(e12))) | ~ (h24(op1(e10, e11)) = op2(h24(e10), h24(e11))) | ~ (h24(op1(e10, e10)) = op2(h24(e10), h24(e10)))) & ((~ (e25 = h23(e15)) & ~ (e25 = h23(e14)) & ~ (e25 = h23(e13)) & ~ (e25 = h23(e12)) & ~ (e25 = h23(e11)) & ~ (e25 = h23(e10))) | sP114 | sP113 | sP112 | sP111 | sP110 | ~ (h23(op1(e15, e15)) = op2(h23(e15), h23(e15))) | ~ (h23(op1(e15, e14)) = op2(h23(e15), h23(e14))) | ~ (h23(op1(e15, e13)) = op2(h23(e15), h23(e13))) | ~ (h23(op1(e15, e12)) = op2(h23(e15), h23(e12))) | ~ (h23(op1(e15, e11)) = op2(h23(e15), h23(e11))) | ~ (h23(op1(e15, e10)) = op2(h23(e15), h23(e10))) | ~ (h23(op1(e14, e15)) = op2(h23(e14), h23(e15))) | ~ (h23(op1(e14, e14)) = op2(h23(e14), h23(e14))) | ~ (h23(op1(e14, e13)) = op2(h23(e14), h23(e13))) | ~ (h23(op1(e14, e12)) = op2(h23(e14), h23(e12))) | ~ (h23(op1(e14, e11)) = op2(h23(e14), h23(e11))) | ~ (h23(op1(e14, e10)) = op2(h23(e14), h23(e10))) | ~ (h23(op1(e13, e15)) = op2(h23(e13), h23(e15))) | ~ (h23(op1(e13, e14)) = op2(h23(e13), h23(e14))) | ~ (h23(op1(e13, e13)) = op2(h23(e13), h23(e13))) | ~ (h23(op1(e13, e12)) = op2(h23(e13), h23(e12))) | ~ (h23(op1(e13, e11)) = op2(h23(e13), h23(e11))) | ~ (h23(op1(e13, e10)) = op2(h23(e13), h23(e10))) | ~ (h23(op1(e12, e15)) = op2(h23(e12), h23(e15))) | ~ (h23(op1(e12, e14)) = op2(h23(e12), h23(e14))) | ~ (h23(op1(e12, e13)) = op2(h23(e12), h23(e13))) | ~ (h23(op1(e12, e12)) = op2(h23(e12), h23(e12))) | ~ (h23(op1(e12, e11)) = op2(h23(e12), h23(e11))) | ~ (h23(op1(e12, e10)) = op2(h23(e12), h23(e10))) | ~ (h23(op1(e11, e15)) = op2(h23(e11), h23(e15))) | ~ (h23(op1(e11, e14)) = op2(h23(e11), h23(e14))) | ~ (h23(op1(e11, e13)) = op2(h23(e11), h23(e13))) | ~ (h23(op1(e11, e12)) = op2(h23(e11), h23(e12))) | ~ (h23(op1(e11, e11)) = op2(h23(e11), h23(e11))) | ~ (h23(op1(e11, e10)) = op2(h23(e11), h23(e10))) | ~ (h23(op1(e10, e15)) = op2(h23(e10), h23(e15))) | ~ (h23(op1(e10, e14)) = op2(h23(e10), h23(e14))) | ~ (h23(op1(e10, e13)) = op2(h23(e10), h23(e13))) | ~ (h23(op1(e10, e12)) = op2(h23(e10), h23(e12))) | ~ (h23(op1(e10, e11)) = op2(h23(e10), h23(e11))) | ~ (h23(op1(e10, e10)) = op2(h23(e10), h23(e10)))) & ((~ (e25 = h22(e15)) & ~ (e25 = h22(e14)) & ~ (e25 = h22(e13)) & ~ (e25 = h22(e12)) & ~ (e25 = h22(e11)) & ~ (e25 = h22(e10))) | sP109 | sP108 | sP107 | sP106 | sP105 | ~ (h22(op1(e15, e15)) = op2(h22(e15), h22(e15))) | ~ (h22(op1(e15, e14)) = op2(h22(e15), h22(e14))) | ~ (h22(op1(e15, e13)) = op2(h22(e15), h22(e13))) | ~ (h22(op1(e15, e12)) = op2(h22(e15), h22(e12))) | ~ (h22(op1(e15, e11)) = op2(h22(e15), h22(e11))) | ~ (h22(op1(e15, e10)) = op2(h22(e15), h22(e10))) | ~ (h22(op1(e14, e15)) = op2(h22(e14), h22(e15))) | ~ (h22(op1(e14, e14)) = op2(h22(e14), h22(e14))) | ~ (h22(op1(e14, e13)) = op2(h22(e14), h22(e13))) | ~ (h22(op1(e14, e12)) = op2(h22(e14), h22(e12))) | ~ (h22(op1(e14, e11)) = op2(h22(e14), h22(e11))) | ~ (h22(op1(e14, e10)) = op2(h22(e14), h22(e10))) | ~ (h22(op1(e13, e15)) = op2(h22(e13), h22(e15))) | ~ (h22(op1(e13, e14)) = op2(h22(e13), h22(e14))) | ~ (h22(op1(e13, e13)) = op2(h22(e13), h22(e13))) | ~ (h22(op1(e13, e12)) = op2(h22(e13), h22(e12))) | ~ (h22(op1(e13, e11)) = op2(h22(e13), h22(e11))) | ~ (h22(op1(e13, e10)) = op2(h22(e13), h22(e10))) | ~ (h22(op1(e12, e15)) = op2(h22(e12), h22(e15))) | ~ (h22(op1(e12, e14)) = op2(h22(e12), h22(e14))) | ~ (h22(op1(e12, e13)) = op2(h22(e12), h22(e13))) | ~ (h22(op1(e12, e12)) = op2(h22(e12), h22(e12))) | ~ (h22(op1(e12, e11)) = op2(h22(e12), h22(e11))) | ~ (h22(op1(e12, e10)) = op2(h22(e12), h22(e10))) | ~ (h22(op1(e11, e15)) = op2(h22(e11), h22(e15))) | ~ (h22(op1(e11, e14)) = op2(h22(e11), h22(e14))) | ~ (h22(op1(e11, e13)) = op2(h22(e11), h22(e13))) | ~ (h22(op1(e11, e12)) = op2(h22(e11), h22(e12))) | ~ (h22(op1(e11, e11)) = op2(h22(e11), h22(e11))) | ~ (h22(op1(e11, e10)) = op2(h22(e11), h22(e10))) | ~ (h22(op1(e10, e15)) = op2(h22(e10), h22(e15))) | ~ (h22(op1(e10, e14)) = op2(h22(e10), h22(e14))) | ~ (h22(op1(e10, e13)) = op2(h22(e10), h22(e13))) | ~ (h22(op1(e10, e12)) = op2(h22(e10), h22(e12))) | ~ (h22(op1(e10, e11)) = op2(h22(e10), h22(e11))) | ~ (h22(op1(e10, e10)) = op2(h22(e10), h22(e10)))) & ((~ (e25 = h21(e15)) & ~ (e25 = h21(e14)) & ~ (e25 = h21(e13)) & ~ (e25 = h21(e12)) & ~ (e25 = h21(e11)) & ~ (e25 = h21(e10))) | sP104 | sP103 | sP102 | sP101 | sP100 | ~ (h21(op1(e15, e15)) = op2(h21(e15), h21(e15))) | ~ (h21(op1(e15, e14)) = op2(h21(e15), h21(e14))) | ~ (h21(op1(e15, e13)) = op2(h21(e15), h21(e13))) | ~ (h21(op1(e15, e12)) = op2(h21(e15), h21(e12))) | ~ (h21(op1(e15, e11)) = op2(h21(e15), h21(e11))) | ~ (h21(op1(e15, e10)) = op2(h21(e15), h21(e10))) | ~ (h21(op1(e14, e15)) = op2(h21(e14), h21(e15))) | ~ (h21(op1(e14, e14)) = op2(h21(e14), h21(e14))) | ~ (h21(op1(e14, e13)) = op2(h21(e14), h21(e13))) | ~ (h21(op1(e14, e12)) = op2(h21(e14), h21(e12))) | ~ (h21(op1(e14, e11)) = op2(h21(e14), h21(e11))) | ~ (h21(op1(e14, e10)) = op2(h21(e14), h21(e10))) | ~ (h21(op1(e13, e15)) = op2(h21(e13), h21(e15))) | ~ (h21(op1(e13, e14)) = op2(h21(e13), h21(e14))) | ~ (h21(op1(e13, e13)) = op2(h21(e13), h21(e13))) | ~ (h21(op1(e13, e12)) = op2(h21(e13), h21(e12))) | ~ (h21(op1(e13, e11)) = op2(h21(e13), h21(e11))) | ~ (h21(op1(e13, e10)) = op2(h21(e13), h21(e10))) | ~ (h21(op1(e12, e15)) = op2(h21(e12), h21(e15))) | ~ (h21(op1(e12, e14)) = op2(h21(e12), h21(e14))) | ~ (h21(op1(e12, e13)) = op2(h21(e12), h21(e13))) | ~ (h21(op1(e12, e12)) = op2(h21(e12), h21(e12))) | ~ (h21(op1(e12, e11)) = op2(h21(e12), h21(e11))) | ~ (h21(op1(e12, e10)) = op2(h21(e12), h21(e10))) | ~ (h21(op1(e11, e15)) = op2(h21(e11), h21(e15))) | ~ (h21(op1(e11, e14)) = op2(h21(e11), h21(e14))) | ~ (h21(op1(e11, e13)) = op2(h21(e11), h21(e13))) | ~ (h21(op1(e11, e12)) = op2(h21(e11), h21(e12))) | ~ (h21(op1(e11, e11)) = op2(h21(e11), h21(e11))) | ~ (h21(op1(e11, e10)) = op2(h21(e11), h21(e10))) | ~ (h21(op1(e10, e15)) = op2(h21(e10), h21(e15))) | ~ (h21(op1(e10, e14)) = op2(h21(e10), h21(e14))) | ~ (h21(op1(e10, e13)) = op2(h21(e10), h21(e13))) | ~ (h21(op1(e10, e12)) = op2(h21(e10), h21(e12))) | ~ (h21(op1(e10, e11)) = op2(h21(e10), h21(e11))) | ~ (h21(op1(e10, e10)) = op2(h21(e10), h21(e10)))) & ((~ (e25 = h20(e15)) & ~ (e25 = h20(e14)) & ~ (e25 = h20(e13)) & ~ (e25 = h20(e12)) & ~ (e25 = h20(e11)) & ~ (e25 = h20(e10))) | sP99 | sP98 | sP97 | sP96 | sP95 | ~ (h20(op1(e15, e15)) = op2(h20(e15), h20(e15))) | ~ (h20(op1(e15, e14)) = op2(h20(e15), h20(e14))) | ~ (h20(op1(e15, e13)) = op2(h20(e15), h20(e13))) | ~ (h20(op1(e15, e12)) = op2(h20(e15), h20(e12))) | ~ (h20(op1(e15, e11)) = op2(h20(e15), h20(e11))) | ~ (h20(op1(e15, e10)) = op2(h20(e15), h20(e10))) | ~ (h20(op1(e14, e15)) = op2(h20(e14), h20(e15))) | ~ (h20(op1(e14, e14)) = op2(h20(e14), h20(e14))) | ~ (h20(op1(e14, e13)) = op2(h20(e14), h20(e13))) | ~ (h20(op1(e14, e12)) = op2(h20(e14), h20(e12))) | ~ (h20(op1(e14, e11)) = op2(h20(e14), h20(e11))) | ~ (h20(op1(e14, e10)) = op2(h20(e14), h20(e10))) | ~ (h20(op1(e13, e15)) = op2(h20(e13), h20(e15))) | ~ (h20(op1(e13, e14)) = op2(h20(e13), h20(e14))) | ~ (h20(op1(e13, e13)) = op2(h20(e13), h20(e13))) | ~ (h20(op1(e13, e12)) = op2(h20(e13), h20(e12))) | ~ (h20(op1(e13, e11)) = op2(h20(e13), h20(e11))) | ~ (h20(op1(e13, e10)) = op2(h20(e13), h20(e10))) | ~ (h20(op1(e12, e15)) = op2(h20(e12), h20(e15))) | ~ (h20(op1(e12, e14)) = op2(h20(e12), h20(e14))) | ~ (h20(op1(e12, e13)) = op2(h20(e12), h20(e13))) | ~ (h20(op1(e12, e12)) = op2(h20(e12), h20(e12))) | ~ (h20(op1(e12, e11)) = op2(h20(e12), h20(e11))) | ~ (h20(op1(e12, e10)) = op2(h20(e12), h20(e10))) | ~ (h20(op1(e11, e15)) = op2(h20(e11), h20(e15))) | ~ (h20(op1(e11, e14)) = op2(h20(e11), h20(e14))) | ~ (h20(op1(e11, e13)) = op2(h20(e11), h20(e13))) | ~ (h20(op1(e11, e12)) = op2(h20(e11), h20(e12))) | ~ (h20(op1(e11, e11)) = op2(h20(e11), h20(e11))) | ~ (h20(op1(e11, e10)) = op2(h20(e11), h20(e10))) | ~ (h20(op1(e10, e15)) = op2(h20(e10), h20(e15))) | ~ (h20(op1(e10, e14)) = op2(h20(e10), h20(e14))) | ~ (h20(op1(e10, e13)) = op2(h20(e10), h20(e13))) | ~ (h20(op1(e10, e12)) = op2(h20(e10), h20(e12))) | ~ (h20(op1(e10, e11)) = op2(h20(e10), h20(e11))) | ~ (h20(op1(e10, e10)) = op2(h20(e10), h20(e10)))) & ((~ (e25 = h19(e15)) & ~ (e25 = h19(e14)) & ~ (e25 = h19(e13)) & ~ (e25 = h19(e12)) & ~ (e25 = h19(e11)) & ~ (e25 = h19(e10))) | sP94 | sP93 | sP92 | sP91 | sP90 | ~ (h19(op1(e15, e15)) = op2(h19(e15), h19(e15))) | ~ (h19(op1(e15, e14)) = op2(h19(e15), h19(e14))) | ~ (h19(op1(e15, e13)) = op2(h19(e15), h19(e13))) | ~ (h19(op1(e15, e12)) = op2(h19(e15), h19(e12))) | ~ (h19(op1(e15, e11)) = op2(h19(e15), h19(e11))) | ~ (h19(op1(e15, e10)) = op2(h19(e15), h19(e10))) | ~ (h19(op1(e14, e15)) = op2(h19(e14), h19(e15))) | ~ (h19(op1(e14, e14)) = op2(h19(e14), h19(e14))) | ~ (h19(op1(e14, e13)) = op2(h19(e14), h19(e13))) | ~ (h19(op1(e14, e12)) = op2(h19(e14), h19(e12))) | ~ (h19(op1(e14, e11)) = op2(h19(e14), h19(e11))) | ~ (h19(op1(e14, e10)) = op2(h19(e14), h19(e10))) | ~ (h19(op1(e13, e15)) = op2(h19(e13), h19(e15))) | ~ (h19(op1(e13, e14)) = op2(h19(e13), h19(e14))) | ~ (h19(op1(e13, e13)) = op2(h19(e13), h19(e13))) | ~ (h19(op1(e13, e12)) = op2(h19(e13), h19(e12))) | ~ (h19(op1(e13, e11)) = op2(h19(e13), h19(e11))) | ~ (h19(op1(e13, e10)) = op2(h19(e13), h19(e10))) | ~ (h19(op1(e12, e15)) = op2(h19(e12), h19(e15))) | ~ (h19(op1(e12, e14)) = op2(h19(e12), h19(e14))) | ~ (h19(op1(e12, e13)) = op2(h19(e12), h19(e13))) | ~ (h19(op1(e12, e12)) = op2(h19(e12), h19(e12))) | ~ (h19(op1(e12, e11)) = op2(h19(e12), h19(e11))) | ~ (h19(op1(e12, e10)) = op2(h19(e12), h19(e10))) | ~ (h19(op1(e11, e15)) = op2(h19(e11), h19(e15))) | ~ (h19(op1(e11, e14)) = op2(h19(e11), h19(e14))) | ~ (h19(op1(e11, e13)) = op2(h19(e11), h19(e13))) | ~ (h19(op1(e11, e12)) = op2(h19(e11), h19(e12))) | ~ (h19(op1(e11, e11)) = op2(h19(e11), h19(e11))) | ~ (h19(op1(e11, e10)) = op2(h19(e11), h19(e10))) | ~ (h19(op1(e10, e15)) = op2(h19(e10), h19(e15))) | ~ (h19(op1(e10, e14)) = op2(h19(e10), h19(e14))) | ~ (h19(op1(e10, e13)) = op2(h19(e10), h19(e13))) | ~ (h19(op1(e10, e12)) = op2(h19(e10), h19(e12))) | ~ (h19(op1(e10, e11)) = op2(h19(e10), h19(e11))) | ~ (h19(op1(e10, e10)) = op2(h19(e10), h19(e10)))) & ((~ (e25 = h18(e15)) & ~ (e25 = h18(e14)) & ~ (e25 = h18(e13)) & ~ (e25 = h18(e12)) & ~ (e25 = h18(e11)) & ~ (e25 = h18(e10))) | sP89 | sP88 | sP87 | sP86 | sP85 | ~ (h18(op1(e15, e15)) = op2(h18(e15), h18(e15))) | ~ (h18(op1(e15, e14)) = op2(h18(e15), h18(e14))) | ~ (h18(op1(e15, e13)) = op2(h18(e15), h18(e13))) | ~ (h18(op1(e15, e12)) = op2(h18(e15), h18(e12))) | ~ (h18(op1(e15, e11)) = op2(h18(e15), h18(e11))) | ~ (h18(op1(e15, e10)) = op2(h18(e15), h18(e10))) | ~ (h18(op1(e14, e15)) = op2(h18(e14), h18(e15))) | ~ (h18(op1(e14, e14)) = op2(h18(e14), h18(e14))) | ~ (h18(op1(e14, e13)) = op2(h18(e14), h18(e13))) | ~ (h18(op1(e14, e12)) = op2(h18(e14), h18(e12))) | ~ (h18(op1(e14, e11)) = op2(h18(e14), h18(e11))) | ~ (h18(op1(e14, e10)) = op2(h18(e14), h18(e10))) | ~ (h18(op1(e13, e15)) = op2(h18(e13), h18(e15))) | ~ (h18(op1(e13, e14)) = op2(h18(e13), h18(e14))) | ~ (h18(op1(e13, e13)) = op2(h18(e13), h18(e13))) | ~ (h18(op1(e13, e12)) = op2(h18(e13), h18(e12))) | ~ (h18(op1(e13, e11)) = op2(h18(e13), h18(e11))) | ~ (h18(op1(e13, e10)) = op2(h18(e13), h18(e10))) | ~ (h18(op1(e12, e15)) = op2(h18(e12), h18(e15))) | ~ (h18(op1(e12, e14)) = op2(h18(e12), h18(e14))) | ~ (h18(op1(e12, e13)) = op2(h18(e12), h18(e13))) | ~ (h18(op1(e12, e12)) = op2(h18(e12), h18(e12))) | ~ (h18(op1(e12, e11)) = op2(h18(e12), h18(e11))) | ~ (h18(op1(e12, e10)) = op2(h18(e12), h18(e10))) | ~ (h18(op1(e11, e15)) = op2(h18(e11), h18(e15))) | ~ (h18(op1(e11, e14)) = op2(h18(e11), h18(e14))) | ~ (h18(op1(e11, e13)) = op2(h18(e11), h18(e13))) | ~ (h18(op1(e11, e12)) = op2(h18(e11), h18(e12))) | ~ (h18(op1(e11, e11)) = op2(h18(e11), h18(e11))) | ~ (h18(op1(e11, e10)) = op2(h18(e11), h18(e10))) | ~ (h18(op1(e10, e15)) = op2(h18(e10), h18(e15))) | ~ (h18(op1(e10, e14)) = op2(h18(e10), h18(e14))) | ~ (h18(op1(e10, e13)) = op2(h18(e10), h18(e13))) | ~ (h18(op1(e10, e12)) = op2(h18(e10), h18(e12))) | ~ (h18(op1(e10, e11)) = op2(h18(e10), h18(e11))) | ~ (h18(op1(e10, e10)) = op2(h18(e10), h18(e10)))) & ((~ (e25 = h17(e15)) & ~ (e25 = h17(e14)) & ~ (e25 = h17(e13)) & ~ (e25 = h17(e12)) & ~ (e25 = h17(e11)) & ~ (e25 = h17(e10))) | sP84 | sP83 | sP82 | sP81 | sP80 | ~ (h17(op1(e15, e15)) = op2(h17(e15), h17(e15))) | ~ (h17(op1(e15, e14)) = op2(h17(e15), h17(e14))) | ~ (h17(op1(e15, e13)) = op2(h17(e15), h17(e13))) | ~ (h17(op1(e15, e12)) = op2(h17(e15), h17(e12))) | ~ (h17(op1(e15, e11)) = op2(h17(e15), h17(e11))) | ~ (h17(op1(e15, e10)) = op2(h17(e15), h17(e10))) | ~ (h17(op1(e14, e15)) = op2(h17(e14), h17(e15))) | ~ (h17(op1(e14, e14)) = op2(h17(e14), h17(e14))) | ~ (h17(op1(e14, e13)) = op2(h17(e14), h17(e13))) | ~ (h17(op1(e14, e12)) = op2(h17(e14), h17(e12))) | ~ (h17(op1(e14, e11)) = op2(h17(e14), h17(e11))) | ~ (h17(op1(e14, e10)) = op2(h17(e14), h17(e10))) | ~ (h17(op1(e13, e15)) = op2(h17(e13), h17(e15))) | ~ (h17(op1(e13, e14)) = op2(h17(e13), h17(e14))) | ~ (h17(op1(e13, e13)) = op2(h17(e13), h17(e13))) | ~ (h17(op1(e13, e12)) = op2(h17(e13), h17(e12))) | ~ (h17(op1(e13, e11)) = op2(h17(e13), h17(e11))) | ~ (h17(op1(e13, e10)) = op2(h17(e13), h17(e10))) | ~ (h17(op1(e12, e15)) = op2(h17(e12), h17(e15))) | ~ (h17(op1(e12, e14)) = op2(h17(e12), h17(e14))) | ~ (h17(op1(e12, e13)) = op2(h17(e12), h17(e13))) | ~ (h17(op1(e12, e12)) = op2(h17(e12), h17(e12))) | ~ (h17(op1(e12, e11)) = op2(h17(e12), h17(e11))) | ~ (h17(op1(e12, e10)) = op2(h17(e12), h17(e10))) | ~ (h17(op1(e11, e15)) = op2(h17(e11), h17(e15))) | ~ (h17(op1(e11, e14)) = op2(h17(e11), h17(e14))) | ~ (h17(op1(e11, e13)) = op2(h17(e11), h17(e13))) | ~ (h17(op1(e11, e12)) = op2(h17(e11), h17(e12))) | ~ (h17(op1(e11, e11)) = op2(h17(e11), h17(e11))) | ~ (h17(op1(e11, e10)) = op2(h17(e11), h17(e10))) | ~ (h17(op1(e10, e15)) = op2(h17(e10), h17(e15))) | ~ (h17(op1(e10, e14)) = op2(h17(e10), h17(e14))) | ~ (h17(op1(e10, e13)) = op2(h17(e10), h17(e13))) | ~ (h17(op1(e10, e12)) = op2(h17(e10), h17(e12))) | ~ (h17(op1(e10, e11)) = op2(h17(e10), h17(e11))) | ~ (h17(op1(e10, e10)) = op2(h17(e10), h17(e10)))) & ((~ (e25 = h16(e15)) & ~ (e25 = h16(e14)) & ~ (e25 = h16(e13)) & ~ (e25 = h16(e12)) & ~ (e25 = h16(e11)) & ~ (e25 = h16(e10))) | sP79 | sP78 | sP77 | sP76 | sP75 | ~ (h16(op1(e15, e15)) = op2(h16(e15), h16(e15))) | ~ (h16(op1(e15, e14)) = op2(h16(e15), h16(e14))) | ~ (h16(op1(e15, e13)) = op2(h16(e15), h16(e13))) | ~ (h16(op1(e15, e12)) = op2(h16(e15), h16(e12))) | ~ (h16(op1(e15, e11)) = op2(h16(e15), h16(e11))) | ~ (h16(op1(e15, e10)) = op2(h16(e15), h16(e10))) | ~ (h16(op1(e14, e15)) = op2(h16(e14), h16(e15))) | ~ (h16(op1(e14, e14)) = op2(h16(e14), h16(e14))) | ~ (h16(op1(e14, e13)) = op2(h16(e14), h16(e13))) | ~ (h16(op1(e14, e12)) = op2(h16(e14), h16(e12))) | ~ (h16(op1(e14, e11)) = op2(h16(e14), h16(e11))) | ~ (h16(op1(e14, e10)) = op2(h16(e14), h16(e10))) | ~ (h16(op1(e13, e15)) = op2(h16(e13), h16(e15))) | ~ (h16(op1(e13, e14)) = op2(h16(e13), h16(e14))) | ~ (h16(op1(e13, e13)) = op2(h16(e13), h16(e13))) | ~ (h16(op1(e13, e12)) = op2(h16(e13), h16(e12))) | ~ (h16(op1(e13, e11)) = op2(h16(e13), h16(e11))) | ~ (h16(op1(e13, e10)) = op2(h16(e13), h16(e10))) | ~ (h16(op1(e12, e15)) = op2(h16(e12), h16(e15))) | ~ (h16(op1(e12, e14)) = op2(h16(e12), h16(e14))) | ~ (h16(op1(e12, e13)) = op2(h16(e12), h16(e13))) | ~ (h16(op1(e12, e12)) = op2(h16(e12), h16(e12))) | ~ (h16(op1(e12, e11)) = op2(h16(e12), h16(e11))) | ~ (h16(op1(e12, e10)) = op2(h16(e12), h16(e10))) | ~ (h16(op1(e11, e15)) = op2(h16(e11), h16(e15))) | ~ (h16(op1(e11, e14)) = op2(h16(e11), h16(e14))) | ~ (h16(op1(e11, e13)) = op2(h16(e11), h16(e13))) | ~ (h16(op1(e11, e12)) = op2(h16(e11), h16(e12))) | ~ (h16(op1(e11, e11)) = op2(h16(e11), h16(e11))) | ~ (h16(op1(e11, e10)) = op2(h16(e11), h16(e10))) | ~ (h16(op1(e10, e15)) = op2(h16(e10), h16(e15))) | ~ (h16(op1(e10, e14)) = op2(h16(e10), h16(e14))) | ~ (h16(op1(e10, e13)) = op2(h16(e10), h16(e13))) | ~ (h16(op1(e10, e12)) = op2(h16(e10), h16(e12))) | ~ (h16(op1(e10, e11)) = op2(h16(e10), h16(e11))) | ~ (h16(op1(e10, e10)) = op2(h16(e10), h16(e10)))) & ((~ (e25 = h15(e15)) & ~ (e25 = h15(e14)) & ~ (e25 = h15(e13)) & ~ (e25 = h15(e12)) & ~ (e25 = h15(e11)) & ~ (e25 = h15(e10))) | sP74 | sP73 | sP72 | sP71 | sP70 | ~ (h15(op1(e15, e15)) = op2(h15(e15), h15(e15))) | ~ (h15(op1(e15, e14)) = op2(h15(e15), h15(e14))) | ~ (h15(op1(e15, e13)) = op2(h15(e15), h15(e13))) | ~ (h15(op1(e15, e12)) = op2(h15(e15), h15(e12))) | ~ (h15(op1(e15, e11)) = op2(h15(e15), h15(e11))) | ~ (h15(op1(e15, e10)) = op2(h15(e15), h15(e10))) | ~ (h15(op1(e14, e15)) = op2(h15(e14), h15(e15))) | ~ (h15(op1(e14, e14)) = op2(h15(e14), h15(e14))) | ~ (h15(op1(e14, e13)) = op2(h15(e14), h15(e13))) | ~ (h15(op1(e14, e12)) = op2(h15(e14), h15(e12))) | ~ (h15(op1(e14, e11)) = op2(h15(e14), h15(e11))) | ~ (h15(op1(e14, e10)) = op2(h15(e14), h15(e10))) | ~ (h15(op1(e13, e15)) = op2(h15(e13), h15(e15))) | ~ (h15(op1(e13, e14)) = op2(h15(e13), h15(e14))) | ~ (h15(op1(e13, e13)) = op2(h15(e13), h15(e13))) | ~ (h15(op1(e13, e12)) = op2(h15(e13), h15(e12))) | ~ (h15(op1(e13, e11)) = op2(h15(e13), h15(e11))) | ~ (h15(op1(e13, e10)) = op2(h15(e13), h15(e10))) | ~ (h15(op1(e12, e15)) = op2(h15(e12), h15(e15))) | ~ (h15(op1(e12, e14)) = op2(h15(e12), h15(e14))) | ~ (h15(op1(e12, e13)) = op2(h15(e12), h15(e13))) | ~ (h15(op1(e12, e12)) = op2(h15(e12), h15(e12))) | ~ (h15(op1(e12, e11)) = op2(h15(e12), h15(e11))) | ~ (h15(op1(e12, e10)) = op2(h15(e12), h15(e10))) | ~ (h15(op1(e11, e15)) = op2(h15(e11), h15(e15))) | ~ (h15(op1(e11, e14)) = op2(h15(e11), h15(e14))) | ~ (h15(op1(e11, e13)) = op2(h15(e11), h15(e13))) | ~ (h15(op1(e11, e12)) = op2(h15(e11), h15(e12))) | ~ (h15(op1(e11, e11)) = op2(h15(e11), h15(e11))) | ~ (h15(op1(e11, e10)) = op2(h15(e11), h15(e10))) | ~ (h15(op1(e10, e15)) = op2(h15(e10), h15(e15))) | ~ (h15(op1(e10, e14)) = op2(h15(e10), h15(e14))) | ~ (h15(op1(e10, e13)) = op2(h15(e10), h15(e13))) | ~ (h15(op1(e10, e12)) = op2(h15(e10), h15(e12))) | ~ (h15(op1(e10, e11)) = op2(h15(e10), h15(e11))) | ~ (h15(op1(e10, e10)) = op2(h15(e10), h15(e10)))) & ((~ (e25 = h14(e15)) & ~ (e25 = h14(e14)) & ~ (e25 = h14(e13)) & ~ (e25 = h14(e12)) & ~ (e25 = h14(e11)) & ~ (e25 = h14(e10))) | sP69 | sP68 | sP67 | sP66 | sP65 | ~ (h14(op1(e15, e15)) = op2(h14(e15), h14(e15))) | ~ (h14(op1(e15, e14)) = op2(h14(e15), h14(e14))) | ~ (h14(op1(e15, e13)) = op2(h14(e15), h14(e13))) | ~ (h14(op1(e15, e12)) = op2(h14(e15), h14(e12))) | ~ (h14(op1(e15, e11)) = op2(h14(e15), h14(e11))) | ~ (h14(op1(e15, e10)) = op2(h14(e15), h14(e10))) | ~ (h14(op1(e14, e15)) = op2(h14(e14), h14(e15))) | ~ (h14(op1(e14, e14)) = op2(h14(e14), h14(e14))) | ~ (h14(op1(e14, e13)) = op2(h14(e14), h14(e13))) | ~ (h14(op1(e14, e12)) = op2(h14(e14), h14(e12))) | ~ (h14(op1(e14, e11)) = op2(h14(e14), h14(e11))) | ~ (h14(op1(e14, e10)) = op2(h14(e14), h14(e10))) | ~ (h14(op1(e13, e15)) = op2(h14(e13), h14(e15))) | ~ (h14(op1(e13, e14)) = op2(h14(e13), h14(e14))) | ~ (h14(op1(e13, e13)) = op2(h14(e13), h14(e13))) | ~ (h14(op1(e13, e12)) = op2(h14(e13), h14(e12))) | ~ (h14(op1(e13, e11)) = op2(h14(e13), h14(e11))) | ~ (h14(op1(e13, e10)) = op2(h14(e13), h14(e10))) | ~ (h14(op1(e12, e15)) = op2(h14(e12), h14(e15))) | ~ (h14(op1(e12, e14)) = op2(h14(e12), h14(e14))) | ~ (h14(op1(e12, e13)) = op2(h14(e12), h14(e13))) | ~ (h14(op1(e12, e12)) = op2(h14(e12), h14(e12))) | ~ (h14(op1(e12, e11)) = op2(h14(e12), h14(e11))) | ~ (h14(op1(e12, e10)) = op2(h14(e12), h14(e10))) | ~ (h14(op1(e11, e15)) = op2(h14(e11), h14(e15))) | ~ (h14(op1(e11, e14)) = op2(h14(e11), h14(e14))) | ~ (h14(op1(e11, e13)) = op2(h14(e11), h14(e13))) | ~ (h14(op1(e11, e12)) = op2(h14(e11), h14(e12))) | ~ (h14(op1(e11, e11)) = op2(h14(e11), h14(e11))) | ~ (h14(op1(e11, e10)) = op2(h14(e11), h14(e10))) | ~ (h14(op1(e10, e15)) = op2(h14(e10), h14(e15))) | ~ (h14(op1(e10, e14)) = op2(h14(e10), h14(e14))) | ~ (h14(op1(e10, e13)) = op2(h14(e10), h14(e13))) | ~ (h14(op1(e10, e12)) = op2(h14(e10), h14(e12))) | ~ (h14(op1(e10, e11)) = op2(h14(e10), h14(e11))) | ~ (h14(op1(e10, e10)) = op2(h14(e10), h14(e10)))) & ((~ (e25 = h13(e15)) & ~ (e25 = h13(e14)) & ~ (e25 = h13(e13)) & ~ (e25 = h13(e12)) & ~ (e25 = h13(e11)) & ~ (e25 = h13(e10))) | sP64 | sP63 | sP62 | sP61 | sP60 | ~ (h13(op1(e15, e15)) = op2(h13(e15), h13(e15))) | ~ (h13(op1(e15, e14)) = op2(h13(e15), h13(e14))) | ~ (h13(op1(e15, e13)) = op2(h13(e15), h13(e13))) | ~ (h13(op1(e15, e12)) = op2(h13(e15), h13(e12))) | ~ (h13(op1(e15, e11)) = op2(h13(e15), h13(e11))) | ~ (h13(op1(e15, e10)) = op2(h13(e15), h13(e10))) | ~ (h13(op1(e14, e15)) = op2(h13(e14), h13(e15))) | ~ (h13(op1(e14, e14)) = op2(h13(e14), h13(e14))) | ~ (h13(op1(e14, e13)) = op2(h13(e14), h13(e13))) | ~ (h13(op1(e14, e12)) = op2(h13(e14), h13(e12))) | ~ (h13(op1(e14, e11)) = op2(h13(e14), h13(e11))) | ~ (h13(op1(e14, e10)) = op2(h13(e14), h13(e10))) | ~ (h13(op1(e13, e15)) = op2(h13(e13), h13(e15))) | ~ (h13(op1(e13, e14)) = op2(h13(e13), h13(e14))) | ~ (h13(op1(e13, e13)) = op2(h13(e13), h13(e13))) | ~ (h13(op1(e13, e12)) = op2(h13(e13), h13(e12))) | ~ (h13(op1(e13, e11)) = op2(h13(e13), h13(e11))) | ~ (h13(op1(e13, e10)) = op2(h13(e13), h13(e10))) | ~ (h13(op1(e12, e15)) = op2(h13(e12), h13(e15))) | ~ (h13(op1(e12, e14)) = op2(h13(e12), h13(e14))) | ~ (h13(op1(e12, e13)) = op2(h13(e12), h13(e13))) | ~ (h13(op1(e12, e12)) = op2(h13(e12), h13(e12))) | ~ (h13(op1(e12, e11)) = op2(h13(e12), h13(e11))) | ~ (h13(op1(e12, e10)) = op2(h13(e12), h13(e10))) | ~ (h13(op1(e11, e15)) = op2(h13(e11), h13(e15))) | ~ (h13(op1(e11, e14)) = op2(h13(e11), h13(e14))) | ~ (h13(op1(e11, e13)) = op2(h13(e11), h13(e13))) | ~ (h13(op1(e11, e12)) = op2(h13(e11), h13(e12))) | ~ (h13(op1(e11, e11)) = op2(h13(e11), h13(e11))) | ~ (h13(op1(e11, e10)) = op2(h13(e11), h13(e10))) | ~ (h13(op1(e10, e15)) = op2(h13(e10), h13(e15))) | ~ (h13(op1(e10, e14)) = op2(h13(e10), h13(e14))) | ~ (h13(op1(e10, e13)) = op2(h13(e10), h13(e13))) | ~ (h13(op1(e10, e12)) = op2(h13(e10), h13(e12))) | ~ (h13(op1(e10, e11)) = op2(h13(e10), h13(e11))) | ~ (h13(op1(e10, e10)) = op2(h13(e10), h13(e10)))) & ((~ (e25 = h12(e15)) & ~ (e25 = h12(e14)) & ~ (e25 = h12(e13)) & ~ (e25 = h12(e12)) & ~ (e25 = h12(e11)) & ~ (e25 = h12(e10))) | sP59 | sP58 | sP57 | sP56 | sP55 | ~ (h12(op1(e15, e15)) = op2(h12(e15), h12(e15))) | ~ (h12(op1(e15, e14)) = op2(h12(e15), h12(e14))) | ~ (h12(op1(e15, e13)) = op2(h12(e15), h12(e13))) | ~ (h12(op1(e15, e12)) = op2(h12(e15), h12(e12))) | ~ (h12(op1(e15, e11)) = op2(h12(e15), h12(e11))) | ~ (h12(op1(e15, e10)) = op2(h12(e15), h12(e10))) | ~ (h12(op1(e14, e15)) = op2(h12(e14), h12(e15))) | ~ (h12(op1(e14, e14)) = op2(h12(e14), h12(e14))) | ~ (h12(op1(e14, e13)) = op2(h12(e14), h12(e13))) | ~ (h12(op1(e14, e12)) = op2(h12(e14), h12(e12))) | ~ (h12(op1(e14, e11)) = op2(h12(e14), h12(e11))) | ~ (h12(op1(e14, e10)) = op2(h12(e14), h12(e10))) | ~ (h12(op1(e13, e15)) = op2(h12(e13), h12(e15))) | ~ (h12(op1(e13, e14)) = op2(h12(e13), h12(e14))) | ~ (h12(op1(e13, e13)) = op2(h12(e13), h12(e13))) | ~ (h12(op1(e13, e12)) = op2(h12(e13), h12(e12))) | ~ (h12(op1(e13, e11)) = op2(h12(e13), h12(e11))) | ~ (h12(op1(e13, e10)) = op2(h12(e13), h12(e10))) | ~ (h12(op1(e12, e15)) = op2(h12(e12), h12(e15))) | ~ (h12(op1(e12, e14)) = op2(h12(e12), h12(e14))) | ~ (h12(op1(e12, e13)) = op2(h12(e12), h12(e13))) | ~ (h12(op1(e12, e12)) = op2(h12(e12), h12(e12))) | ~ (h12(op1(e12, e11)) = op2(h12(e12), h12(e11))) | ~ (h12(op1(e12, e10)) = op2(h12(e12), h12(e10))) | ~ (h12(op1(e11, e15)) = op2(h12(e11), h12(e15))) | ~ (h12(op1(e11, e14)) = op2(h12(e11), h12(e14))) | ~ (h12(op1(e11, e13)) = op2(h12(e11), h12(e13))) | ~ (h12(op1(e11, e12)) = op2(h12(e11), h12(e12))) | ~ (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) | ~ (h12(op1(e11, e10)) = op2(h12(e11), h12(e10))) | ~ (h12(op1(e10, e15)) = op2(h12(e10), h12(e15))) | ~ (h12(op1(e10, e14)) = op2(h12(e10), h12(e14))) | ~ (h12(op1(e10, e13)) = op2(h12(e10), h12(e13))) | ~ (h12(op1(e10, e12)) = op2(h12(e10), h12(e12))) | ~ (h12(op1(e10, e11)) = op2(h12(e10), h12(e11))) | ~ (h12(op1(e10, e10)) = op2(h12(e10), h12(e10)))) & ((~ (e25 = h11(e15)) & ~ (e25 = h11(e14)) & ~ (e25 = h11(e13)) & ~ (e25 = h11(e12)) & ~ (e25 = h11(e11)) & ~ (e25 = h11(e10))) | sP54 | sP53 | sP52 | sP51 | sP50 | ~ (h11(op1(e15, e15)) = op2(h11(e15), h11(e15))) | ~ (h11(op1(e15, e14)) = op2(h11(e15), h11(e14))) | ~ (h11(op1(e15, e13)) = op2(h11(e15), h11(e13))) | ~ (h11(op1(e15, e12)) = op2(h11(e15), h11(e12))) | ~ (h11(op1(e15, e11)) = op2(h11(e15), h11(e11))) | ~ (h11(op1(e15, e10)) = op2(h11(e15), h11(e10))) | ~ (h11(op1(e14, e15)) = op2(h11(e14), h11(e15))) | ~ (h11(op1(e14, e14)) = op2(h11(e14), h11(e14))) | ~ (h11(op1(e14, e13)) = op2(h11(e14), h11(e13))) | ~ (h11(op1(e14, e12)) = op2(h11(e14), h11(e12))) | ~ (h11(op1(e14, e11)) = op2(h11(e14), h11(e11))) | ~ (h11(op1(e14, e10)) = op2(h11(e14), h11(e10))) | ~ (h11(op1(e13, e15)) = op2(h11(e13), h11(e15))) | ~ (h11(op1(e13, e14)) = op2(h11(e13), h11(e14))) | ~ (h11(op1(e13, e13)) = op2(h11(e13), h11(e13))) | ~ (h11(op1(e13, e12)) = op2(h11(e13), h11(e12))) | ~ (h11(op1(e13, e11)) = op2(h11(e13), h11(e11))) | ~ (h11(op1(e13, e10)) = op2(h11(e13), h11(e10))) | ~ (h11(op1(e12, e15)) = op2(h11(e12), h11(e15))) | ~ (h11(op1(e12, e14)) = op2(h11(e12), h11(e14))) | ~ (h11(op1(e12, e13)) = op2(h11(e12), h11(e13))) | ~ (h11(op1(e12, e12)) = op2(h11(e12), h11(e12))) | ~ (h11(op1(e12, e11)) = op2(h11(e12), h11(e11))) | ~ (h11(op1(e12, e10)) = op2(h11(e12), h11(e10))) | ~ (h11(op1(e11, e15)) = op2(h11(e11), h11(e15))) | ~ (h11(op1(e11, e14)) = op2(h11(e11), h11(e14))) | ~ (h11(op1(e11, e13)) = op2(h11(e11), h11(e13))) | ~ (h11(op1(e11, e12)) = op2(h11(e11), h11(e12))) | ~ (h11(op1(e11, e11)) = op2(h11(e11), h11(e11))) | ~ (h11(op1(e11, e10)) = op2(h11(e11), h11(e10))) | ~ (h11(op1(e10, e15)) = op2(h11(e10), h11(e15))) | ~ (h11(op1(e10, e14)) = op2(h11(e10), h11(e14))) | ~ (h11(op1(e10, e13)) = op2(h11(e10), h11(e13))) | ~ (h11(op1(e10, e12)) = op2(h11(e10), h11(e12))) | ~ (h11(op1(e10, e11)) = op2(h11(e10), h11(e11))) | ~ (h11(op1(e10, e10)) = op2(h11(e10), h11(e10)))) & ((~ (e25 = h10(e15)) & ~ (e25 = h10(e14)) & ~ (e25 = h10(e13)) & ~ (e25 = h10(e12)) & ~ (e25 = h10(e11)) & ~ (e25 = h10(e10))) | sP49 | sP48 | sP47 | sP46 | sP45 | ~ (h10(op1(e15, e15)) = op2(h10(e15), h10(e15))) | ~ (h10(op1(e15, e14)) = op2(h10(e15), h10(e14))) | ~ (h10(op1(e15, e13)) = op2(h10(e15), h10(e13))) | ~ (h10(op1(e15, e12)) = op2(h10(e15), h10(e12))) | ~ (h10(op1(e15, e11)) = op2(h10(e15), h10(e11))) | ~ (h10(op1(e15, e10)) = op2(h10(e15), h10(e10))) | ~ (h10(op1(e14, e15)) = op2(h10(e14), h10(e15))) | ~ (h10(op1(e14, e14)) = op2(h10(e14), h10(e14))) | ~ (h10(op1(e14, e13)) = op2(h10(e14), h10(e13))) | ~ (h10(op1(e14, e12)) = op2(h10(e14), h10(e12))) | ~ (h10(op1(e14, e11)) = op2(h10(e14), h10(e11))) | ~ (h10(op1(e14, e10)) = op2(h10(e14), h10(e10))) | ~ (h10(op1(e13, e15)) = op2(h10(e13), h10(e15))) | ~ (h10(op1(e13, e14)) = op2(h10(e13), h10(e14))) | ~ (h10(op1(e13, e13)) = op2(h10(e13), h10(e13))) | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) | ~ (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) | ~ (h10(op1(e12, e15)) = op2(h10(e12), h10(e15))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), h10(e14))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) | ~ (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) | ~ (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) | ~ (h10(op1(e11, e15)) = op2(h10(e11), h10(e15))) | ~ (h10(op1(e11, e14)) = op2(h10(e11), h10(e14))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e15)) = op2(h10(e10), h10(e15))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))) & ((~ (e25 = h9(e15)) & ~ (e25 = h9(e14)) & ~ (e25 = h9(e13)) & ~ (e25 = h9(e12)) & ~ (e25 = h9(e11)) & ~ (e25 = h9(e10))) | sP44 | sP43 | sP42 | sP41 | sP40 | ~ (h9(op1(e15, e15)) = op2(h9(e15), h9(e15))) | ~ (h9(op1(e15, e14)) = op2(h9(e15), h9(e14))) | ~ (h9(op1(e15, e13)) = op2(h9(e15), h9(e13))) | ~ (h9(op1(e15, e12)) = op2(h9(e15), h9(e12))) | ~ (h9(op1(e15, e11)) = op2(h9(e15), h9(e11))) | ~ (h9(op1(e15, e10)) = op2(h9(e15), h9(e10))) | ~ (h9(op1(e14, e15)) = op2(h9(e14), h9(e15))) | ~ (h9(op1(e14, e14)) = op2(h9(e14), h9(e14))) | ~ (h9(op1(e14, e13)) = op2(h9(e14), h9(e13))) | ~ (h9(op1(e14, e12)) = op2(h9(e14), h9(e12))) | ~ (h9(op1(e14, e11)) = op2(h9(e14), h9(e11))) | ~ (h9(op1(e14, e10)) = op2(h9(e14), h9(e10))) | ~ (h9(op1(e13, e15)) = op2(h9(e13), h9(e15))) | ~ (h9(op1(e13, e14)) = op2(h9(e13), h9(e14))) | ~ (h9(op1(e13, e13)) = op2(h9(e13), h9(e13))) | ~ (h9(op1(e13, e12)) = op2(h9(e13), h9(e12))) | ~ (h9(op1(e13, e11)) = op2(h9(e13), h9(e11))) | ~ (h9(op1(e13, e10)) = op2(h9(e13), h9(e10))) | ~ (h9(op1(e12, e15)) = op2(h9(e12), h9(e15))) | ~ (h9(op1(e12, e14)) = op2(h9(e12), h9(e14))) | ~ (h9(op1(e12, e13)) = op2(h9(e12), h9(e13))) | ~ (h9(op1(e12, e12)) = op2(h9(e12), h9(e12))) | ~ (h9(op1(e12, e11)) = op2(h9(e12), h9(e11))) | ~ (h9(op1(e12, e10)) = op2(h9(e12), h9(e10))) | ~ (h9(op1(e11, e15)) = op2(h9(e11), h9(e15))) | ~ (h9(op1(e11, e14)) = op2(h9(e11), h9(e14))) | ~ (h9(op1(e11, e13)) = op2(h9(e11), h9(e13))) | ~ (h9(op1(e11, e12)) = op2(h9(e11), h9(e12))) | ~ (h9(op1(e11, e11)) = op2(h9(e11), h9(e11))) | ~ (h9(op1(e11, e10)) = op2(h9(e11), h9(e10))) | ~ (h9(op1(e10, e15)) = op2(h9(e10), h9(e15))) | ~ (h9(op1(e10, e14)) = op2(h9(e10), h9(e14))) | ~ (h9(op1(e10, e13)) = op2(h9(e10), h9(e13))) | ~ (h9(op1(e10, e12)) = op2(h9(e10), h9(e12))) | ~ (h9(op1(e10, e11)) = op2(h9(e10), h9(e11))) | ~ (h9(op1(e10, e10)) = op2(h9(e10), h9(e10)))) & ((~ (e25 = h8(e15)) & ~ (e25 = h8(e14)) & ~ (e25 = h8(e13)) & ~ (e25 = h8(e12)) & ~ (e25 = h8(e11)) & ~ (e25 = h8(e10))) | sP39 | sP38 | sP37 | sP36 | sP35 | ~ (h8(op1(e15, e15)) = op2(h8(e15), h8(e15))) | ~ (h8(op1(e15, e14)) = op2(h8(e15), h8(e14))) | ~ (h8(op1(e15, e13)) = op2(h8(e15), h8(e13))) | ~ (h8(op1(e15, e12)) = op2(h8(e15), h8(e12))) | ~ (h8(op1(e15, e11)) = op2(h8(e15), h8(e11))) | ~ (h8(op1(e15, e10)) = op2(h8(e15), h8(e10))) | ~ (h8(op1(e14, e15)) = op2(h8(e14), h8(e15))) | ~ (h8(op1(e14, e14)) = op2(h8(e14), h8(e14))) | ~ (h8(op1(e14, e13)) = op2(h8(e14), h8(e13))) | ~ (h8(op1(e14, e12)) = op2(h8(e14), h8(e12))) | ~ (h8(op1(e14, e11)) = op2(h8(e14), h8(e11))) | ~ (h8(op1(e14, e10)) = op2(h8(e14), h8(e10))) | ~ (h8(op1(e13, e15)) = op2(h8(e13), h8(e15))) | ~ (h8(op1(e13, e14)) = op2(h8(e13), h8(e14))) | ~ (h8(op1(e13, e13)) = op2(h8(e13), h8(e13))) | ~ (h8(op1(e13, e12)) = op2(h8(e13), h8(e12))) | ~ (h8(op1(e13, e11)) = op2(h8(e13), h8(e11))) | ~ (h8(op1(e13, e10)) = op2(h8(e13), h8(e10))) | ~ (h8(op1(e12, e15)) = op2(h8(e12), h8(e15))) | ~ (h8(op1(e12, e14)) = op2(h8(e12), h8(e14))) | ~ (h8(op1(e12, e13)) = op2(h8(e12), h8(e13))) | ~ (h8(op1(e12, e12)) = op2(h8(e12), h8(e12))) | ~ (h8(op1(e12, e11)) = op2(h8(e12), h8(e11))) | ~ (h8(op1(e12, e10)) = op2(h8(e12), h8(e10))) | ~ (h8(op1(e11, e15)) = op2(h8(e11), h8(e15))) | ~ (h8(op1(e11, e14)) = op2(h8(e11), h8(e14))) | ~ (h8(op1(e11, e13)) = op2(h8(e11), h8(e13))) | ~ (h8(op1(e11, e12)) = op2(h8(e11), h8(e12))) | ~ (h8(op1(e11, e11)) = op2(h8(e11), h8(e11))) | ~ (h8(op1(e11, e10)) = op2(h8(e11), h8(e10))) | ~ (h8(op1(e10, e15)) = op2(h8(e10), h8(e15))) | ~ (h8(op1(e10, e14)) = op2(h8(e10), h8(e14))) | ~ (h8(op1(e10, e13)) = op2(h8(e10), h8(e13))) | ~ (h8(op1(e10, e12)) = op2(h8(e10), h8(e12))) | ~ (h8(op1(e10, e11)) = op2(h8(e10), h8(e11))) | ~ (h8(op1(e10, e10)) = op2(h8(e10), h8(e10)))) & ((~ (e25 = h7(e15)) & ~ (e25 = h7(e14)) & ~ (e25 = h7(e13)) & ~ (e25 = h7(e12)) & ~ (e25 = h7(e11)) & ~ (e25 = h7(e10))) | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h7(op1(e15, e15)) = op2(h7(e15), h7(e15))) | ~ (h7(op1(e15, e14)) = op2(h7(e15), h7(e14))) | ~ (h7(op1(e15, e13)) = op2(h7(e15), h7(e13))) | ~ (h7(op1(e15, e12)) = op2(h7(e15), h7(e12))) | ~ (h7(op1(e15, e11)) = op2(h7(e15), h7(e11))) | ~ (h7(op1(e15, e10)) = op2(h7(e15), h7(e10))) | ~ (h7(op1(e14, e15)) = op2(h7(e14), h7(e15))) | ~ (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) | ~ (h7(op1(e14, e13)) = op2(h7(e14), h7(e13))) | ~ (h7(op1(e14, e12)) = op2(h7(e14), h7(e12))) | ~ (h7(op1(e14, e11)) = op2(h7(e14), h7(e11))) | ~ (h7(op1(e14, e10)) = op2(h7(e14), h7(e10))) | ~ (h7(op1(e13, e15)) = op2(h7(e13), h7(e15))) | ~ (h7(op1(e13, e14)) = op2(h7(e13), h7(e14))) | ~ (h7(op1(e13, e13)) = op2(h7(e13), h7(e13))) | ~ (h7(op1(e13, e12)) = op2(h7(e13), h7(e12))) | ~ (h7(op1(e13, e11)) = op2(h7(e13), h7(e11))) | ~ (h7(op1(e13, e10)) = op2(h7(e13), h7(e10))) | ~ (h7(op1(e12, e15)) = op2(h7(e12), h7(e15))) | ~ (h7(op1(e12, e14)) = op2(h7(e12), h7(e14))) | ~ (h7(op1(e12, e13)) = op2(h7(e12), h7(e13))) | ~ (h7(op1(e12, e12)) = op2(h7(e12), h7(e12))) | ~ (h7(op1(e12, e11)) = op2(h7(e12), h7(e11))) | ~ (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) | ~ (h7(op1(e11, e15)) = op2(h7(e11), h7(e15))) | ~ (h7(op1(e11, e14)) = op2(h7(e11), h7(e14))) | ~ (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) | ~ (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) | ~ (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) | ~ (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) | ~ (h7(op1(e10, e15)) = op2(h7(e10), h7(e15))) | ~ (h7(op1(e10, e14)) = op2(h7(e10), h7(e14))) | ~ (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) | ~ (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) | ~ (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))) & ((~ (e25 = h6(e15)) & ~ (e25 = h6(e14)) & ~ (e25 = h6(e13)) & ~ (e25 = h6(e12)) & ~ (e25 = h6(e11)) & ~ (e25 = h6(e10))) | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (h6(op1(e15, e15)) = op2(h6(e15), h6(e15))) | ~ (h6(op1(e15, e14)) = op2(h6(e15), h6(e14))) | ~ (h6(op1(e15, e13)) = op2(h6(e15), h6(e13))) | ~ (h6(op1(e15, e12)) = op2(h6(e15), h6(e12))) | ~ (h6(op1(e15, e11)) = op2(h6(e15), h6(e11))) | ~ (h6(op1(e15, e10)) = op2(h6(e15), h6(e10))) | ~ (h6(op1(e14, e15)) = op2(h6(e14), h6(e15))) | ~ (h6(op1(e14, e14)) = op2(h6(e14), h6(e14))) | ~ (h6(op1(e14, e13)) = op2(h6(e14), h6(e13))) | ~ (h6(op1(e14, e12)) = op2(h6(e14), h6(e12))) | ~ (h6(op1(e14, e11)) = op2(h6(e14), h6(e11))) | ~ (h6(op1(e14, e10)) = op2(h6(e14), h6(e10))) | ~ (h6(op1(e13, e15)) = op2(h6(e13), h6(e15))) | ~ (h6(op1(e13, e14)) = op2(h6(e13), h6(e14))) | ~ (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) | ~ (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e15)) = op2(h6(e12), h6(e15))) | ~ (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))) & ((~ (e25 = h5(e15)) & ~ (e25 = h5(e14)) & ~ (e25 = h5(e13)) & ~ (e25 = h5(e12)) & ~ (e25 = h5(e11)) & ~ (e25 = h5(e10))) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e15, e14)) = op2(h5(e15), h5(e14))) | ~ (h5(op1(e15, e13)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e15, e12)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e15, e11)) = op2(h5(e15), h5(e11))) | ~ (h5(op1(e15, e10)) = op2(h5(e15), h5(e10))) | ~ (h5(op1(e14, e15)) = op2(h5(e14), h5(e15))) | ~ (h5(op1(e14, e14)) = op2(h5(e14), h5(e14))) | ~ (h5(op1(e14, e13)) = op2(h5(e14), h5(e13))) | ~ (h5(op1(e14, e12)) = op2(h5(e14), h5(e12))) | ~ (h5(op1(e14, e11)) = op2(h5(e14), h5(e11))) | ~ (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e15)) = op2(h5(e11), h5(e15))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))) & ((~ (e25 = h4(e15)) & ~ (e25 = h4(e14)) & ~ (e25 = h4(e13)) & ~ (e25 = h4(e12)) & ~ (e25 = h4(e11)) & ~ (e25 = h4(e10))) | sP19 | sP18 | sP17 | sP16 | sP15 | ~ (h4(op1(e15, e15)) = op2(h4(e15), h4(e15))) | ~ (h4(op1(e15, e14)) = op2(h4(e15), h4(e14))) | ~ (h4(op1(e15, e13)) = op2(h4(e15), h4(e13))) | ~ (h4(op1(e15, e12)) = op2(h4(e15), h4(e12))) | ~ (h4(op1(e15, e11)) = op2(h4(e15), h4(e11))) | ~ (h4(op1(e15, e10)) = op2(h4(e15), h4(e10))) | ~ (h4(op1(e14, e15)) = op2(h4(e14), h4(e15))) | ~ (h4(op1(e14, e14)) = op2(h4(e14), h4(e14))) | ~ (h4(op1(e14, e13)) = op2(h4(e14), h4(e13))) | ~ (h4(op1(e14, e12)) = op2(h4(e14), h4(e12))) | ~ (h4(op1(e14, e11)) = op2(h4(e14), h4(e11))) | ~ (h4(op1(e14, e10)) = op2(h4(e14), h4(e10))) | ~ (h4(op1(e13, e15)) = op2(h4(e13), h4(e15))) | ~ (h4(op1(e13, e14)) = op2(h4(e13), h4(e14))) | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e15)) = op2(h4(e12), h4(e15))) | ~ (h4(op1(e12, e14)) = op2(h4(e12), h4(e14))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e15)) = op2(h4(e11), h4(e15))) | ~ (h4(op1(e11, e14)) = op2(h4(e11), h4(e14))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e15)) = op2(h4(e10), h4(e15))) | ~ (h4(op1(e10, e14)) = op2(h4(e10), h4(e14))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) & ((~ (e25 = h3(e15)) & ~ (e25 = h3(e14)) & ~ (e25 = h3(e13)) & ~ (e25 = h3(e12)) & ~ (e25 = h3(e11)) & ~ (e25 = h3(e10))) | sP14 | sP13 | sP12 | sP11 | sP10 | ~ (h3(op1(e15, e15)) = op2(h3(e15), h3(e15))) | ~ (h3(op1(e15, e14)) = op2(h3(e15), h3(e14))) | ~ (h3(op1(e15, e13)) = op2(h3(e15), h3(e13))) | ~ (h3(op1(e15, e12)) = op2(h3(e15), h3(e12))) | ~ (h3(op1(e15, e11)) = op2(h3(e15), h3(e11))) | ~ (h3(op1(e15, e10)) = op2(h3(e15), h3(e10))) | ~ (h3(op1(e14, e15)) = op2(h3(e14), h3(e15))) | ~ (h3(op1(e14, e14)) = op2(h3(e14), h3(e14))) | ~ (h3(op1(e14, e13)) = op2(h3(e14), h3(e13))) | ~ (h3(op1(e14, e12)) = op2(h3(e14), h3(e12))) | ~ (h3(op1(e14, e11)) = op2(h3(e14), h3(e11))) | ~ (h3(op1(e14, e10)) = op2(h3(e14), h3(e10))) | ~ (h3(op1(e13, e15)) = op2(h3(e13), h3(e15))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), h3(e14))) | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e15)) = op2(h3(e12), h3(e15))) | ~ (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e15)) = op2(h3(e11), h3(e15))) | ~ (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e15)) = op2(h3(e10), h3(e15))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) & ((~ (e25 = h2(e15)) & ~ (e25 = h2(e14)) & ~ (e25 = h2(e13)) & ~ (e25 = h2(e12)) & ~ (e25 = h2(e11)) & ~ (e25 = h2(e10))) | sP9 | sP8 | sP7 | sP6 | sP5 | ~ (h2(op1(e15, e15)) = op2(h2(e15), h2(e15))) | ~ (h2(op1(e15, e14)) = op2(h2(e15), h2(e14))) | ~ (h2(op1(e15, e13)) = op2(h2(e15), h2(e13))) | ~ (h2(op1(e15, e12)) = op2(h2(e15), h2(e12))) | ~ (h2(op1(e15, e11)) = op2(h2(e15), h2(e11))) | ~ (h2(op1(e15, e10)) = op2(h2(e15), h2(e10))) | ~ (h2(op1(e14, e15)) = op2(h2(e14), h2(e15))) | ~ (h2(op1(e14, e14)) = op2(h2(e14), h2(e14))) | ~ (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) | ~ (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) | ~ (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) | ~ (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) | ~ (h2(op1(e13, e15)) = op2(h2(e13), h2(e15))) | ~ (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e15)) = op2(h2(e12), h2(e15))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e15)) = op2(h2(e11), h2(e15))) | ~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e15)) = op2(h2(e10), h2(e15))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) & ((~ (e25 = h1(e15)) & ~ (e25 = h1(e14)) & ~ (e25 = h1(e13)) & ~ (e25 = h1(e12)) & ~ (e25 = h1(e11)) & ~ (e25 = h1(e10))) | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (h1(op1(e15, e15)) = op2(h1(e15), h1(e15))) | ~ (h1(op1(e15, e14)) = op2(h1(e15), h1(e14))) | ~ (h1(op1(e15, e13)) = op2(h1(e15), h1(e13))) | ~ (h1(op1(e15, e12)) = op2(h1(e15), h1(e12))) | ~ (h1(op1(e15, e11)) = op2(h1(e15), h1(e11))) | ~ (h1(op1(e15, e10)) = op2(h1(e15), h1(e10))) | ~ (h1(op1(e14, e15)) = op2(h1(e14), h1(e15))) | ~ (h1(op1(e14, e14)) = op2(h1(e14), h1(e14))) | ~ (h1(op1(e14, e13)) = op2(h1(e14), h1(e13))) | ~ (h1(op1(e14, e12)) = op2(h1(e14), h1(e12))) | ~ (h1(op1(e14, e11)) = op2(h1(e14), h1(e11))) | ~ (h1(op1(e14, e10)) = op2(h1(e14), h1(e10))) | ~ (h1(op1(e13, e15)) = op2(h1(e13), h1(e15))) | ~ (h1(op1(e13, e14)) = op2(h1(e13), h1(e14))) | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e15)) = op2(h1(e12), h1(e15))) | ~ (h1(op1(e12, e14)) = op2(h1(e12), h1(e14))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e15)) = op2(h1(e11), h1(e15))) | ~ (h1(op1(e11, e14)) = op2(h1(e11), h1(e14))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e15)) = op2(h1(e10), h1(e15))) | ~ (h1(op1(e10, e14)) = op2(h1(e10), h1(e14))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(definition_folding, [], [f62, e212, e211, e210, e209, e208, e207, e206, e205, e204, e203, e202, e201, e200, e199, e198, e197, e196, e195, e194, e193, e192, e191, e190, e189, e188, e187, e186, e185, e184, e183, e182, e181, e180, e179, e178, e177, e176, e175, e174, e173, e172, e171, e170, e169, e168, e167, e166, e165, e164, e163, e162, e161, e160, e159, e158, e157, e156, e155, e154, e153, e152, e151, e150, e149, e148, e147, e146, e145, e144, e143, e142, e141, e140, e139, e138, e137, e136, e135, e134, e133, e132, e131, e130, e129, e128, e127, e126, e125, e124, e123, e122, e121, e120, e119, e118, e117, e116, e115, e114, e113, e112, e111, e110, e109, e108, e107, e106, e105, e104, e103, e102, e101, e100, e99, e98, e97, e96, e95, e94, e93, e92, e91, e90, e89, e88, e87, e86, e85, e84, e83, e82, e81, e80, e79, e78, e77, e76, e75, e74, e73, e72, e71, e70, e69, e68, e67, e66, e65, e64, e63])).
fof(f63, plain, ((~ (e20 = h1(e15)) & ~ (e20 = h1(e14)) & ~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10))) | ~ sP0), inference(usedef, [], [e63])).
fof(e63, plain, (sP0 <=> (~ (e20 = h1(e15)) & ~ (e20 = h1(e14)) & ~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f64, plain, ((~ (e21 = h1(e15)) & ~ (e21 = h1(e14)) & ~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10))) | ~ sP1), inference(usedef, [], [e64])).
fof(e64, plain, (sP1 <=> (~ (e21 = h1(e15)) & ~ (e21 = h1(e14)) & ~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f65, plain, ((~ (e22 = h1(e15)) & ~ (e22 = h1(e14)) & ~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10))) | ~ sP2), inference(usedef, [], [e65])).
fof(e65, plain, (sP2 <=> (~ (e22 = h1(e15)) & ~ (e22 = h1(e14)) & ~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f66, plain, ((~ (e23 = h1(e15)) & ~ (e23 = h1(e14)) & ~ (e23 = h1(e13)) & ~ (e23 = h1(e12)) & ~ (e23 = h1(e11)) & ~ (e23 = h1(e10))) | ~ sP3), inference(usedef, [], [e66])).
fof(e66, plain, (sP3 <=> (~ (e23 = h1(e15)) & ~ (e23 = h1(e14)) & ~ (e23 = h1(e13)) & ~ (e23 = h1(e12)) & ~ (e23 = h1(e11)) & ~ (e23 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f67, plain, ((~ (e24 = h1(e15)) & ~ (e24 = h1(e14)) & ~ (e24 = h1(e13)) & ~ (e24 = h1(e12)) & ~ (e24 = h1(e11)) & ~ (e24 = h1(e10))) | ~ sP4), inference(usedef, [], [e67])).
fof(e67, plain, (sP4 <=> (~ (e24 = h1(e15)) & ~ (e24 = h1(e14)) & ~ (e24 = h1(e13)) & ~ (e24 = h1(e12)) & ~ (e24 = h1(e11)) & ~ (e24 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f68, plain, ((~ (e20 = h2(e15)) & ~ (e20 = h2(e14)) & ~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10))) | ~ sP5), inference(usedef, [], [e68])).
fof(e68, plain, (sP5 <=> (~ (e20 = h2(e15)) & ~ (e20 = h2(e14)) & ~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f69, plain, ((~ (e21 = h2(e15)) & ~ (e21 = h2(e14)) & ~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10))) | ~ sP6), inference(usedef, [], [e69])).
fof(e69, plain, (sP6 <=> (~ (e21 = h2(e15)) & ~ (e21 = h2(e14)) & ~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f70, plain, ((~ (e22 = h2(e15)) & ~ (e22 = h2(e14)) & ~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10))) | ~ sP7), inference(usedef, [], [e70])).
fof(e70, plain, (sP7 <=> (~ (e22 = h2(e15)) & ~ (e22 = h2(e14)) & ~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f71, plain, ((~ (e23 = h2(e15)) & ~ (e23 = h2(e14)) & ~ (e23 = h2(e13)) & ~ (e23 = h2(e12)) & ~ (e23 = h2(e11)) & ~ (e23 = h2(e10))) | ~ sP8), inference(usedef, [], [e71])).
fof(e71, plain, (sP8 <=> (~ (e23 = h2(e15)) & ~ (e23 = h2(e14)) & ~ (e23 = h2(e13)) & ~ (e23 = h2(e12)) & ~ (e23 = h2(e11)) & ~ (e23 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f72, plain, ((~ (e24 = h2(e15)) & ~ (e24 = h2(e14)) & ~ (e24 = h2(e13)) & ~ (e24 = h2(e12)) & ~ (e24 = h2(e11)) & ~ (e24 = h2(e10))) | ~ sP9), inference(usedef, [], [e72])).
fof(e72, plain, (sP9 <=> (~ (e24 = h2(e15)) & ~ (e24 = h2(e14)) & ~ (e24 = h2(e13)) & ~ (e24 = h2(e12)) & ~ (e24 = h2(e11)) & ~ (e24 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f73, plain, ((~ (e20 = h3(e15)) & ~ (e20 = h3(e14)) & ~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10))) | ~ sP10), inference(usedef, [], [e73])).
fof(e73, plain, (sP10 <=> (~ (e20 = h3(e15)) & ~ (e20 = h3(e14)) & ~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f74, plain, ((~ (e21 = h3(e15)) & ~ (e21 = h3(e14)) & ~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10))) | ~ sP11), inference(usedef, [], [e74])).
fof(e74, plain, (sP11 <=> (~ (e21 = h3(e15)) & ~ (e21 = h3(e14)) & ~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f75, plain, ((~ (e22 = h3(e15)) & ~ (e22 = h3(e14)) & ~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10))) | ~ sP12), inference(usedef, [], [e75])).
fof(e75, plain, (sP12 <=> (~ (e22 = h3(e15)) & ~ (e22 = h3(e14)) & ~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f76, plain, ((~ (e23 = h3(e15)) & ~ (e23 = h3(e14)) & ~ (e23 = h3(e13)) & ~ (e23 = h3(e12)) & ~ (e23 = h3(e11)) & ~ (e23 = h3(e10))) | ~ sP13), inference(usedef, [], [e76])).
fof(e76, plain, (sP13 <=> (~ (e23 = h3(e15)) & ~ (e23 = h3(e14)) & ~ (e23 = h3(e13)) & ~ (e23 = h3(e12)) & ~ (e23 = h3(e11)) & ~ (e23 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f77, plain, ((~ (e24 = h3(e15)) & ~ (e24 = h3(e14)) & ~ (e24 = h3(e13)) & ~ (e24 = h3(e12)) & ~ (e24 = h3(e11)) & ~ (e24 = h3(e10))) | ~ sP14), inference(usedef, [], [e77])).
fof(e77, plain, (sP14 <=> (~ (e24 = h3(e15)) & ~ (e24 = h3(e14)) & ~ (e24 = h3(e13)) & ~ (e24 = h3(e12)) & ~ (e24 = h3(e11)) & ~ (e24 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f78, plain, ((~ (e20 = h4(e15)) & ~ (e20 = h4(e14)) & ~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10))) | ~ sP15), inference(usedef, [], [e78])).
fof(e78, plain, (sP15 <=> (~ (e20 = h4(e15)) & ~ (e20 = h4(e14)) & ~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f79, plain, ((~ (e21 = h4(e15)) & ~ (e21 = h4(e14)) & ~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10))) | ~ sP16), inference(usedef, [], [e79])).
fof(e79, plain, (sP16 <=> (~ (e21 = h4(e15)) & ~ (e21 = h4(e14)) & ~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f80, plain, ((~ (e22 = h4(e15)) & ~ (e22 = h4(e14)) & ~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10))) | ~ sP17), inference(usedef, [], [e80])).
fof(e80, plain, (sP17 <=> (~ (e22 = h4(e15)) & ~ (e22 = h4(e14)) & ~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f81, plain, ((~ (e23 = h4(e15)) & ~ (e23 = h4(e14)) & ~ (e23 = h4(e13)) & ~ (e23 = h4(e12)) & ~ (e23 = h4(e11)) & ~ (e23 = h4(e10))) | ~ sP18), inference(usedef, [], [e81])).
fof(e81, plain, (sP18 <=> (~ (e23 = h4(e15)) & ~ (e23 = h4(e14)) & ~ (e23 = h4(e13)) & ~ (e23 = h4(e12)) & ~ (e23 = h4(e11)) & ~ (e23 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f82, plain, ((~ (e24 = h4(e15)) & ~ (e24 = h4(e14)) & ~ (e24 = h4(e13)) & ~ (e24 = h4(e12)) & ~ (e24 = h4(e11)) & ~ (e24 = h4(e10))) | ~ sP19), inference(usedef, [], [e82])).
fof(e82, plain, (sP19 <=> (~ (e24 = h4(e15)) & ~ (e24 = h4(e14)) & ~ (e24 = h4(e13)) & ~ (e24 = h4(e12)) & ~ (e24 = h4(e11)) & ~ (e24 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP19])])).
fof(f83, plain, ((~ (e20 = h5(e15)) & ~ (e20 = h5(e14)) & ~ (e20 = h5(e13)) & ~ (e20 = h5(e12)) & ~ (e20 = h5(e11)) & ~ (e20 = h5(e10))) | ~ sP20), inference(usedef, [], [e83])).
fof(e83, plain, (sP20 <=> (~ (e20 = h5(e15)) & ~ (e20 = h5(e14)) & ~ (e20 = h5(e13)) & ~ (e20 = h5(e12)) & ~ (e20 = h5(e11)) & ~ (e20 = h5(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP20])])).
fof(f84, plain, ((~ (e21 = h5(e15)) & ~ (e21 = h5(e14)) & ~ (e21 = h5(e13)) & ~ (e21 = h5(e12)) & ~ (e21 = h5(e11)) & ~ (e21 = h5(e10))) | ~ sP21), inference(usedef, [], [e84])).
fof(e84, plain, (sP21 <=> (~ (e21 = h5(e15)) & ~ (e21 = h5(e14)) & ~ (e21 = h5(e13)) & ~ (e21 = h5(e12)) & ~ (e21 = h5(e11)) & ~ (e21 = h5(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP21])])).
fof(f85, plain, ((~ (e22 = h5(e15)) & ~ (e22 = h5(e14)) & ~ (e22 = h5(e13)) & ~ (e22 = h5(e12)) & ~ (e22 = h5(e11)) & ~ (e22 = h5(e10))) | ~ sP22), inference(usedef, [], [e85])).
fof(e85, plain, (sP22 <=> (~ (e22 = h5(e15)) & ~ (e22 = h5(e14)) & ~ (e22 = h5(e13)) & ~ (e22 = h5(e12)) & ~ (e22 = h5(e11)) & ~ (e22 = h5(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP22])])).
fof(f86, plain, ((~ (e23 = h5(e15)) & ~ (e23 = h5(e14)) & ~ (e23 = h5(e13)) & ~ (e23 = h5(e12)) & ~ (e23 = h5(e11)) & ~ (e23 = h5(e10))) | ~ sP23), inference(usedef, [], [e86])).
fof(e86, plain, (sP23 <=> (~ (e23 = h5(e15)) & ~ (e23 = h5(e14)) & ~ (e23 = h5(e13)) & ~ (e23 = h5(e12)) & ~ (e23 = h5(e11)) & ~ (e23 = h5(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP23])])).
fof(f87, plain, ((~ (e24 = h5(e15)) & ~ (e24 = h5(e14)) & ~ (e24 = h5(e13)) & ~ (e24 = h5(e12)) & ~ (e24 = h5(e11)) & ~ (e24 = h5(e10))) | ~ sP24), inference(usedef, [], [e87])).
fof(e87, plain, (sP24 <=> (~ (e24 = h5(e15)) & ~ (e24 = h5(e14)) & ~ (e24 = h5(e13)) & ~ (e24 = h5(e12)) & ~ (e24 = h5(e11)) & ~ (e24 = h5(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP24])])).
fof(f88, plain, ((~ (e20 = h6(e15)) & ~ (e20 = h6(e14)) & ~ (e20 = h6(e13)) & ~ (e20 = h6(e12)) & ~ (e20 = h6(e11)) & ~ (e20 = h6(e10))) | ~ sP25), inference(usedef, [], [e88])).
fof(e88, plain, (sP25 <=> (~ (e20 = h6(e15)) & ~ (e20 = h6(e14)) & ~ (e20 = h6(e13)) & ~ (e20 = h6(e12)) & ~ (e20 = h6(e11)) & ~ (e20 = h6(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP25])])).
fof(f89, plain, ((~ (e21 = h6(e15)) & ~ (e21 = h6(e14)) & ~ (e21 = h6(e13)) & ~ (e21 = h6(e12)) & ~ (e21 = h6(e11)) & ~ (e21 = h6(e10))) | ~ sP26), inference(usedef, [], [e89])).
fof(e89, plain, (sP26 <=> (~ (e21 = h6(e15)) & ~ (e21 = h6(e14)) & ~ (e21 = h6(e13)) & ~ (e21 = h6(e12)) & ~ (e21 = h6(e11)) & ~ (e21 = h6(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP26])])).
fof(f90, plain, ((~ (e22 = h6(e15)) & ~ (e22 = h6(e14)) & ~ (e22 = h6(e13)) & ~ (e22 = h6(e12)) & ~ (e22 = h6(e11)) & ~ (e22 = h6(e10))) | ~ sP27), inference(usedef, [], [e90])).
fof(e90, plain, (sP27 <=> (~ (e22 = h6(e15)) & ~ (e22 = h6(e14)) & ~ (e22 = h6(e13)) & ~ (e22 = h6(e12)) & ~ (e22 = h6(e11)) & ~ (e22 = h6(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP27])])).
fof(f91, plain, ((~ (e23 = h6(e15)) & ~ (e23 = h6(e14)) & ~ (e23 = h6(e13)) & ~ (e23 = h6(e12)) & ~ (e23 = h6(e11)) & ~ (e23 = h6(e10))) | ~ sP28), inference(usedef, [], [e91])).
fof(e91, plain, (sP28 <=> (~ (e23 = h6(e15)) & ~ (e23 = h6(e14)) & ~ (e23 = h6(e13)) & ~ (e23 = h6(e12)) & ~ (e23 = h6(e11)) & ~ (e23 = h6(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP28])])).
fof(f92, plain, ((~ (e24 = h6(e15)) & ~ (e24 = h6(e14)) & ~ (e24 = h6(e13)) & ~ (e24 = h6(e12)) & ~ (e24 = h6(e11)) & ~ (e24 = h6(e10))) | ~ sP29), inference(usedef, [], [e92])).
fof(e92, plain, (sP29 <=> (~ (e24 = h6(e15)) & ~ (e24 = h6(e14)) & ~ (e24 = h6(e13)) & ~ (e24 = h6(e12)) & ~ (e24 = h6(e11)) & ~ (e24 = h6(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP29])])).
fof(f93, plain, ((~ (e20 = h7(e15)) & ~ (e20 = h7(e14)) & ~ (e20 = h7(e13)) & ~ (e20 = h7(e12)) & ~ (e20 = h7(e11)) & ~ (e20 = h7(e10))) | ~ sP30), inference(usedef, [], [e93])).
fof(e93, plain, (sP30 <=> (~ (e20 = h7(e15)) & ~ (e20 = h7(e14)) & ~ (e20 = h7(e13)) & ~ (e20 = h7(e12)) & ~ (e20 = h7(e11)) & ~ (e20 = h7(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP30])])).
fof(f94, plain, ((~ (e21 = h7(e15)) & ~ (e21 = h7(e14)) & ~ (e21 = h7(e13)) & ~ (e21 = h7(e12)) & ~ (e21 = h7(e11)) & ~ (e21 = h7(e10))) | ~ sP31), inference(usedef, [], [e94])).
fof(e94, plain, (sP31 <=> (~ (e21 = h7(e15)) & ~ (e21 = h7(e14)) & ~ (e21 = h7(e13)) & ~ (e21 = h7(e12)) & ~ (e21 = h7(e11)) & ~ (e21 = h7(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP31])])).
fof(f95, plain, ((~ (e22 = h7(e15)) & ~ (e22 = h7(e14)) & ~ (e22 = h7(e13)) & ~ (e22 = h7(e12)) & ~ (e22 = h7(e11)) & ~ (e22 = h7(e10))) | ~ sP32), inference(usedef, [], [e95])).
fof(e95, plain, (sP32 <=> (~ (e22 = h7(e15)) & ~ (e22 = h7(e14)) & ~ (e22 = h7(e13)) & ~ (e22 = h7(e12)) & ~ (e22 = h7(e11)) & ~ (e22 = h7(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP32])])).
fof(f96, plain, ((~ (e23 = h7(e15)) & ~ (e23 = h7(e14)) & ~ (e23 = h7(e13)) & ~ (e23 = h7(e12)) & ~ (e23 = h7(e11)) & ~ (e23 = h7(e10))) | ~ sP33), inference(usedef, [], [e96])).
fof(e96, plain, (sP33 <=> (~ (e23 = h7(e15)) & ~ (e23 = h7(e14)) & ~ (e23 = h7(e13)) & ~ (e23 = h7(e12)) & ~ (e23 = h7(e11)) & ~ (e23 = h7(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP33])])).
fof(f97, plain, ((~ (e24 = h7(e15)) & ~ (e24 = h7(e14)) & ~ (e24 = h7(e13)) & ~ (e24 = h7(e12)) & ~ (e24 = h7(e11)) & ~ (e24 = h7(e10))) | ~ sP34), inference(usedef, [], [e97])).
fof(e97, plain, (sP34 <=> (~ (e24 = h7(e15)) & ~ (e24 = h7(e14)) & ~ (e24 = h7(e13)) & ~ (e24 = h7(e12)) & ~ (e24 = h7(e11)) & ~ (e24 = h7(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP34])])).
fof(f98, plain, ((~ (e20 = h8(e15)) & ~ (e20 = h8(e14)) & ~ (e20 = h8(e13)) & ~ (e20 = h8(e12)) & ~ (e20 = h8(e11)) & ~ (e20 = h8(e10))) | ~ sP35), inference(usedef, [], [e98])).
fof(e98, plain, (sP35 <=> (~ (e20 = h8(e15)) & ~ (e20 = h8(e14)) & ~ (e20 = h8(e13)) & ~ (e20 = h8(e12)) & ~ (e20 = h8(e11)) & ~ (e20 = h8(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP35])])).
fof(f99, plain, ((~ (e21 = h8(e15)) & ~ (e21 = h8(e14)) & ~ (e21 = h8(e13)) & ~ (e21 = h8(e12)) & ~ (e21 = h8(e11)) & ~ (e21 = h8(e10))) | ~ sP36), inference(usedef, [], [e99])).
fof(e99, plain, (sP36 <=> (~ (e21 = h8(e15)) & ~ (e21 = h8(e14)) & ~ (e21 = h8(e13)) & ~ (e21 = h8(e12)) & ~ (e21 = h8(e11)) & ~ (e21 = h8(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP36])])).
fof(f100, plain, ((~ (e22 = h8(e15)) & ~ (e22 = h8(e14)) & ~ (e22 = h8(e13)) & ~ (e22 = h8(e12)) & ~ (e22 = h8(e11)) & ~ (e22 = h8(e10))) | ~ sP37), inference(usedef, [], [e100])).
fof(e100, plain, (sP37 <=> (~ (e22 = h8(e15)) & ~ (e22 = h8(e14)) & ~ (e22 = h8(e13)) & ~ (e22 = h8(e12)) & ~ (e22 = h8(e11)) & ~ (e22 = h8(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP37])])).
fof(f101, plain, ((~ (e23 = h8(e15)) & ~ (e23 = h8(e14)) & ~ (e23 = h8(e13)) & ~ (e23 = h8(e12)) & ~ (e23 = h8(e11)) & ~ (e23 = h8(e10))) | ~ sP38), inference(usedef, [], [e101])).
fof(e101, plain, (sP38 <=> (~ (e23 = h8(e15)) & ~ (e23 = h8(e14)) & ~ (e23 = h8(e13)) & ~ (e23 = h8(e12)) & ~ (e23 = h8(e11)) & ~ (e23 = h8(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP38])])).
fof(f102, plain, ((~ (e24 = h8(e15)) & ~ (e24 = h8(e14)) & ~ (e24 = h8(e13)) & ~ (e24 = h8(e12)) & ~ (e24 = h8(e11)) & ~ (e24 = h8(e10))) | ~ sP39), inference(usedef, [], [e102])).
fof(e102, plain, (sP39 <=> (~ (e24 = h8(e15)) & ~ (e24 = h8(e14)) & ~ (e24 = h8(e13)) & ~ (e24 = h8(e12)) & ~ (e24 = h8(e11)) & ~ (e24 = h8(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP39])])).
fof(f103, plain, ((~ (e20 = h9(e15)) & ~ (e20 = h9(e14)) & ~ (e20 = h9(e13)) & ~ (e20 = h9(e12)) & ~ (e20 = h9(e11)) & ~ (e20 = h9(e10))) | ~ sP40), inference(usedef, [], [e103])).
fof(e103, plain, (sP40 <=> (~ (e20 = h9(e15)) & ~ (e20 = h9(e14)) & ~ (e20 = h9(e13)) & ~ (e20 = h9(e12)) & ~ (e20 = h9(e11)) & ~ (e20 = h9(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP40])])).
fof(f104, plain, ((~ (e21 = h9(e15)) & ~ (e21 = h9(e14)) & ~ (e21 = h9(e13)) & ~ (e21 = h9(e12)) & ~ (e21 = h9(e11)) & ~ (e21 = h9(e10))) | ~ sP41), inference(usedef, [], [e104])).
fof(e104, plain, (sP41 <=> (~ (e21 = h9(e15)) & ~ (e21 = h9(e14)) & ~ (e21 = h9(e13)) & ~ (e21 = h9(e12)) & ~ (e21 = h9(e11)) & ~ (e21 = h9(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP41])])).
fof(f105, plain, ((~ (e22 = h9(e15)) & ~ (e22 = h9(e14)) & ~ (e22 = h9(e13)) & ~ (e22 = h9(e12)) & ~ (e22 = h9(e11)) & ~ (e22 = h9(e10))) | ~ sP42), inference(usedef, [], [e105])).
fof(e105, plain, (sP42 <=> (~ (e22 = h9(e15)) & ~ (e22 = h9(e14)) & ~ (e22 = h9(e13)) & ~ (e22 = h9(e12)) & ~ (e22 = h9(e11)) & ~ (e22 = h9(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP42])])).
fof(f106, plain, ((~ (e23 = h9(e15)) & ~ (e23 = h9(e14)) & ~ (e23 = h9(e13)) & ~ (e23 = h9(e12)) & ~ (e23 = h9(e11)) & ~ (e23 = h9(e10))) | ~ sP43), inference(usedef, [], [e106])).
fof(e106, plain, (sP43 <=> (~ (e23 = h9(e15)) & ~ (e23 = h9(e14)) & ~ (e23 = h9(e13)) & ~ (e23 = h9(e12)) & ~ (e23 = h9(e11)) & ~ (e23 = h9(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP43])])).
fof(f107, plain, ((~ (e24 = h9(e15)) & ~ (e24 = h9(e14)) & ~ (e24 = h9(e13)) & ~ (e24 = h9(e12)) & ~ (e24 = h9(e11)) & ~ (e24 = h9(e10))) | ~ sP44), inference(usedef, [], [e107])).
fof(e107, plain, (sP44 <=> (~ (e24 = h9(e15)) & ~ (e24 = h9(e14)) & ~ (e24 = h9(e13)) & ~ (e24 = h9(e12)) & ~ (e24 = h9(e11)) & ~ (e24 = h9(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP44])])).
fof(f108, plain, ((~ (e20 = h10(e15)) & ~ (e20 = h10(e14)) & ~ (e20 = h10(e13)) & ~ (e20 = h10(e12)) & ~ (e20 = h10(e11)) & ~ (e20 = h10(e10))) | ~ sP45), inference(usedef, [], [e108])).
fof(e108, plain, (sP45 <=> (~ (e20 = h10(e15)) & ~ (e20 = h10(e14)) & ~ (e20 = h10(e13)) & ~ (e20 = h10(e12)) & ~ (e20 = h10(e11)) & ~ (e20 = h10(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP45])])).
fof(f109, plain, ((~ (e21 = h10(e15)) & ~ (e21 = h10(e14)) & ~ (e21 = h10(e13)) & ~ (e21 = h10(e12)) & ~ (e21 = h10(e11)) & ~ (e21 = h10(e10))) | ~ sP46), inference(usedef, [], [e109])).
fof(e109, plain, (sP46 <=> (~ (e21 = h10(e15)) & ~ (e21 = h10(e14)) & ~ (e21 = h10(e13)) & ~ (e21 = h10(e12)) & ~ (e21 = h10(e11)) & ~ (e21 = h10(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP46])])).
fof(f110, plain, ((~ (e22 = h10(e15)) & ~ (e22 = h10(e14)) & ~ (e22 = h10(e13)) & ~ (e22 = h10(e12)) & ~ (e22 = h10(e11)) & ~ (e22 = h10(e10))) | ~ sP47), inference(usedef, [], [e110])).
fof(e110, plain, (sP47 <=> (~ (e22 = h10(e15)) & ~ (e22 = h10(e14)) & ~ (e22 = h10(e13)) & ~ (e22 = h10(e12)) & ~ (e22 = h10(e11)) & ~ (e22 = h10(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP47])])).
fof(f111, plain, ((~ (e23 = h10(e15)) & ~ (e23 = h10(e14)) & ~ (e23 = h10(e13)) & ~ (e23 = h10(e12)) & ~ (e23 = h10(e11)) & ~ (e23 = h10(e10))) | ~ sP48), inference(usedef, [], [e111])).
fof(e111, plain, (sP48 <=> (~ (e23 = h10(e15)) & ~ (e23 = h10(e14)) & ~ (e23 = h10(e13)) & ~ (e23 = h10(e12)) & ~ (e23 = h10(e11)) & ~ (e23 = h10(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP48])])).
fof(f113, plain, ((~ (e20 = h11(e15)) & ~ (e20 = h11(e14)) & ~ (e20 = h11(e13)) & ~ (e20 = h11(e12)) & ~ (e20 = h11(e11)) & ~ (e20 = h11(e10))) | ~ sP50), inference(usedef, [], [e113])).
fof(e113, plain, (sP50 <=> (~ (e20 = h11(e15)) & ~ (e20 = h11(e14)) & ~ (e20 = h11(e13)) & ~ (e20 = h11(e12)) & ~ (e20 = h11(e11)) & ~ (e20 = h11(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP50])])).
fof(f114, plain, ((~ (e21 = h11(e15)) & ~ (e21 = h11(e14)) & ~ (e21 = h11(e13)) & ~ (e21 = h11(e12)) & ~ (e21 = h11(e11)) & ~ (e21 = h11(e10))) | ~ sP51), inference(usedef, [], [e114])).
fof(e114, plain, (sP51 <=> (~ (e21 = h11(e15)) & ~ (e21 = h11(e14)) & ~ (e21 = h11(e13)) & ~ (e21 = h11(e12)) & ~ (e21 = h11(e11)) & ~ (e21 = h11(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP51])])).
fof(f115, plain, ((~ (e22 = h11(e15)) & ~ (e22 = h11(e14)) & ~ (e22 = h11(e13)) & ~ (e22 = h11(e12)) & ~ (e22 = h11(e11)) & ~ (e22 = h11(e10))) | ~ sP52), inference(usedef, [], [e115])).
fof(e115, plain, (sP52 <=> (~ (e22 = h11(e15)) & ~ (e22 = h11(e14)) & ~ (e22 = h11(e13)) & ~ (e22 = h11(e12)) & ~ (e22 = h11(e11)) & ~ (e22 = h11(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP52])])).
fof(f116, plain, ((~ (e23 = h11(e15)) & ~ (e23 = h11(e14)) & ~ (e23 = h11(e13)) & ~ (e23 = h11(e12)) & ~ (e23 = h11(e11)) & ~ (e23 = h11(e10))) | ~ sP53), inference(usedef, [], [e116])).
fof(e116, plain, (sP53 <=> (~ (e23 = h11(e15)) & ~ (e23 = h11(e14)) & ~ (e23 = h11(e13)) & ~ (e23 = h11(e12)) & ~ (e23 = h11(e11)) & ~ (e23 = h11(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP53])])).
fof(f117, plain, ((~ (e24 = h11(e15)) & ~ (e24 = h11(e14)) & ~ (e24 = h11(e13)) & ~ (e24 = h11(e12)) & ~ (e24 = h11(e11)) & ~ (e24 = h11(e10))) | ~ sP54), inference(usedef, [], [e117])).
fof(e117, plain, (sP54 <=> (~ (e24 = h11(e15)) & ~ (e24 = h11(e14)) & ~ (e24 = h11(e13)) & ~ (e24 = h11(e12)) & ~ (e24 = h11(e11)) & ~ (e24 = h11(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP54])])).
fof(f118, plain, ((~ (e20 = h12(e15)) & ~ (e20 = h12(e14)) & ~ (e20 = h12(e13)) & ~ (e20 = h12(e12)) & ~ (e20 = h12(e11)) & ~ (e20 = h12(e10))) | ~ sP55), inference(usedef, [], [e118])).
fof(e118, plain, (sP55 <=> (~ (e20 = h12(e15)) & ~ (e20 = h12(e14)) & ~ (e20 = h12(e13)) & ~ (e20 = h12(e12)) & ~ (e20 = h12(e11)) & ~ (e20 = h12(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP55])])).
fof(f119, plain, ((~ (e21 = h12(e15)) & ~ (e21 = h12(e14)) & ~ (e21 = h12(e13)) & ~ (e21 = h12(e12)) & ~ (e21 = h12(e11)) & ~ (e21 = h12(e10))) | ~ sP56), inference(usedef, [], [e119])).
fof(e119, plain, (sP56 <=> (~ (e21 = h12(e15)) & ~ (e21 = h12(e14)) & ~ (e21 = h12(e13)) & ~ (e21 = h12(e12)) & ~ (e21 = h12(e11)) & ~ (e21 = h12(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP56])])).
fof(f120, plain, ((~ (e22 = h12(e15)) & ~ (e22 = h12(e14)) & ~ (e22 = h12(e13)) & ~ (e22 = h12(e12)) & ~ (e22 = h12(e11)) & ~ (e22 = h12(e10))) | ~ sP57), inference(usedef, [], [e120])).
fof(e120, plain, (sP57 <=> (~ (e22 = h12(e15)) & ~ (e22 = h12(e14)) & ~ (e22 = h12(e13)) & ~ (e22 = h12(e12)) & ~ (e22 = h12(e11)) & ~ (e22 = h12(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP57])])).
fof(f121, plain, ((~ (e23 = h12(e15)) & ~ (e23 = h12(e14)) & ~ (e23 = h12(e13)) & ~ (e23 = h12(e12)) & ~ (e23 = h12(e11)) & ~ (e23 = h12(e10))) | ~ sP58), inference(usedef, [], [e121])).
fof(e121, plain, (sP58 <=> (~ (e23 = h12(e15)) & ~ (e23 = h12(e14)) & ~ (e23 = h12(e13)) & ~ (e23 = h12(e12)) & ~ (e23 = h12(e11)) & ~ (e23 = h12(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP58])])).
fof(f122, plain, ((~ (e24 = h12(e15)) & ~ (e24 = h12(e14)) & ~ (e24 = h12(e13)) & ~ (e24 = h12(e12)) & ~ (e24 = h12(e11)) & ~ (e24 = h12(e10))) | ~ sP59), inference(usedef, [], [e122])).
fof(e122, plain, (sP59 <=> (~ (e24 = h12(e15)) & ~ (e24 = h12(e14)) & ~ (e24 = h12(e13)) & ~ (e24 = h12(e12)) & ~ (e24 = h12(e11)) & ~ (e24 = h12(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP59])])).
fof(f123, plain, ((~ (e20 = h13(e15)) & ~ (e20 = h13(e14)) & ~ (e20 = h13(e13)) & ~ (e20 = h13(e12)) & ~ (e20 = h13(e11)) & ~ (e20 = h13(e10))) | ~ sP60), inference(usedef, [], [e123])).
fof(e123, plain, (sP60 <=> (~ (e20 = h13(e15)) & ~ (e20 = h13(e14)) & ~ (e20 = h13(e13)) & ~ (e20 = h13(e12)) & ~ (e20 = h13(e11)) & ~ (e20 = h13(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP60])])).
fof(f124, plain, ((~ (e21 = h13(e15)) & ~ (e21 = h13(e14)) & ~ (e21 = h13(e13)) & ~ (e21 = h13(e12)) & ~ (e21 = h13(e11)) & ~ (e21 = h13(e10))) | ~ sP61), inference(usedef, [], [e124])).
fof(e124, plain, (sP61 <=> (~ (e21 = h13(e15)) & ~ (e21 = h13(e14)) & ~ (e21 = h13(e13)) & ~ (e21 = h13(e12)) & ~ (e21 = h13(e11)) & ~ (e21 = h13(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP61])])).
fof(f125, plain, ((~ (e22 = h13(e15)) & ~ (e22 = h13(e14)) & ~ (e22 = h13(e13)) & ~ (e22 = h13(e12)) & ~ (e22 = h13(e11)) & ~ (e22 = h13(e10))) | ~ sP62), inference(usedef, [], [e125])).
fof(e125, plain, (sP62 <=> (~ (e22 = h13(e15)) & ~ (e22 = h13(e14)) & ~ (e22 = h13(e13)) & ~ (e22 = h13(e12)) & ~ (e22 = h13(e11)) & ~ (e22 = h13(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP62])])).
fof(f126, plain, ((~ (e23 = h13(e15)) & ~ (e23 = h13(e14)) & ~ (e23 = h13(e13)) & ~ (e23 = h13(e12)) & ~ (e23 = h13(e11)) & ~ (e23 = h13(e10))) | ~ sP63), inference(usedef, [], [e126])).
fof(e126, plain, (sP63 <=> (~ (e23 = h13(e15)) & ~ (e23 = h13(e14)) & ~ (e23 = h13(e13)) & ~ (e23 = h13(e12)) & ~ (e23 = h13(e11)) & ~ (e23 = h13(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP63])])).
fof(f127, plain, ((~ (e24 = h13(e15)) & ~ (e24 = h13(e14)) & ~ (e24 = h13(e13)) & ~ (e24 = h13(e12)) & ~ (e24 = h13(e11)) & ~ (e24 = h13(e10))) | ~ sP64), inference(usedef, [], [e127])).
fof(e127, plain, (sP64 <=> (~ (e24 = h13(e15)) & ~ (e24 = h13(e14)) & ~ (e24 = h13(e13)) & ~ (e24 = h13(e12)) & ~ (e24 = h13(e11)) & ~ (e24 = h13(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP64])])).
fof(f128, plain, ((~ (e20 = h14(e15)) & ~ (e20 = h14(e14)) & ~ (e20 = h14(e13)) & ~ (e20 = h14(e12)) & ~ (e20 = h14(e11)) & ~ (e20 = h14(e10))) | ~ sP65), inference(usedef, [], [e128])).
fof(e128, plain, (sP65 <=> (~ (e20 = h14(e15)) & ~ (e20 = h14(e14)) & ~ (e20 = h14(e13)) & ~ (e20 = h14(e12)) & ~ (e20 = h14(e11)) & ~ (e20 = h14(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP65])])).
fof(f129, plain, ((~ (e21 = h14(e15)) & ~ (e21 = h14(e14)) & ~ (e21 = h14(e13)) & ~ (e21 = h14(e12)) & ~ (e21 = h14(e11)) & ~ (e21 = h14(e10))) | ~ sP66), inference(usedef, [], [e129])).
fof(e129, plain, (sP66 <=> (~ (e21 = h14(e15)) & ~ (e21 = h14(e14)) & ~ (e21 = h14(e13)) & ~ (e21 = h14(e12)) & ~ (e21 = h14(e11)) & ~ (e21 = h14(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP66])])).
fof(f130, plain, ((~ (e22 = h14(e15)) & ~ (e22 = h14(e14)) & ~ (e22 = h14(e13)) & ~ (e22 = h14(e12)) & ~ (e22 = h14(e11)) & ~ (e22 = h14(e10))) | ~ sP67), inference(usedef, [], [e130])).
fof(e130, plain, (sP67 <=> (~ (e22 = h14(e15)) & ~ (e22 = h14(e14)) & ~ (e22 = h14(e13)) & ~ (e22 = h14(e12)) & ~ (e22 = h14(e11)) & ~ (e22 = h14(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP67])])).
fof(f131, plain, ((~ (e23 = h14(e15)) & ~ (e23 = h14(e14)) & ~ (e23 = h14(e13)) & ~ (e23 = h14(e12)) & ~ (e23 = h14(e11)) & ~ (e23 = h14(e10))) | ~ sP68), inference(usedef, [], [e131])).
fof(e131, plain, (sP68 <=> (~ (e23 = h14(e15)) & ~ (e23 = h14(e14)) & ~ (e23 = h14(e13)) & ~ (e23 = h14(e12)) & ~ (e23 = h14(e11)) & ~ (e23 = h14(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP68])])).
fof(f132, plain, ((~ (e24 = h14(e15)) & ~ (e24 = h14(e14)) & ~ (e24 = h14(e13)) & ~ (e24 = h14(e12)) & ~ (e24 = h14(e11)) & ~ (e24 = h14(e10))) | ~ sP69), inference(usedef, [], [e132])).
fof(e132, plain, (sP69 <=> (~ (e24 = h14(e15)) & ~ (e24 = h14(e14)) & ~ (e24 = h14(e13)) & ~ (e24 = h14(e12)) & ~ (e24 = h14(e11)) & ~ (e24 = h14(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP69])])).
fof(f133, plain, ((~ (e20 = h15(e15)) & ~ (e20 = h15(e14)) & ~ (e20 = h15(e13)) & ~ (e20 = h15(e12)) & ~ (e20 = h15(e11)) & ~ (e20 = h15(e10))) | ~ sP70), inference(usedef, [], [e133])).
fof(e133, plain, (sP70 <=> (~ (e20 = h15(e15)) & ~ (e20 = h15(e14)) & ~ (e20 = h15(e13)) & ~ (e20 = h15(e12)) & ~ (e20 = h15(e11)) & ~ (e20 = h15(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP70])])).
fof(f134, plain, ((~ (e21 = h15(e15)) & ~ (e21 = h15(e14)) & ~ (e21 = h15(e13)) & ~ (e21 = h15(e12)) & ~ (e21 = h15(e11)) & ~ (e21 = h15(e10))) | ~ sP71), inference(usedef, [], [e134])).
fof(e134, plain, (sP71 <=> (~ (e21 = h15(e15)) & ~ (e21 = h15(e14)) & ~ (e21 = h15(e13)) & ~ (e21 = h15(e12)) & ~ (e21 = h15(e11)) & ~ (e21 = h15(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP71])])).
fof(f135, plain, ((~ (e22 = h15(e15)) & ~ (e22 = h15(e14)) & ~ (e22 = h15(e13)) & ~ (e22 = h15(e12)) & ~ (e22 = h15(e11)) & ~ (e22 = h15(e10))) | ~ sP72), inference(usedef, [], [e135])).
fof(e135, plain, (sP72 <=> (~ (e22 = h15(e15)) & ~ (e22 = h15(e14)) & ~ (e22 = h15(e13)) & ~ (e22 = h15(e12)) & ~ (e22 = h15(e11)) & ~ (e22 = h15(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP72])])).
fof(f136, plain, ((~ (e23 = h15(e15)) & ~ (e23 = h15(e14)) & ~ (e23 = h15(e13)) & ~ (e23 = h15(e12)) & ~ (e23 = h15(e11)) & ~ (e23 = h15(e10))) | ~ sP73), inference(usedef, [], [e136])).
fof(e136, plain, (sP73 <=> (~ (e23 = h15(e15)) & ~ (e23 = h15(e14)) & ~ (e23 = h15(e13)) & ~ (e23 = h15(e12)) & ~ (e23 = h15(e11)) & ~ (e23 = h15(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP73])])).
fof(f137, plain, ((~ (e24 = h15(e15)) & ~ (e24 = h15(e14)) & ~ (e24 = h15(e13)) & ~ (e24 = h15(e12)) & ~ (e24 = h15(e11)) & ~ (e24 = h15(e10))) | ~ sP74), inference(usedef, [], [e137])).
fof(e137, plain, (sP74 <=> (~ (e24 = h15(e15)) & ~ (e24 = h15(e14)) & ~ (e24 = h15(e13)) & ~ (e24 = h15(e12)) & ~ (e24 = h15(e11)) & ~ (e24 = h15(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP74])])).
fof(f138, plain, ((~ (e20 = h16(e15)) & ~ (e20 = h16(e14)) & ~ (e20 = h16(e13)) & ~ (e20 = h16(e12)) & ~ (e20 = h16(e11)) & ~ (e20 = h16(e10))) | ~ sP75), inference(usedef, [], [e138])).
fof(e138, plain, (sP75 <=> (~ (e20 = h16(e15)) & ~ (e20 = h16(e14)) & ~ (e20 = h16(e13)) & ~ (e20 = h16(e12)) & ~ (e20 = h16(e11)) & ~ (e20 = h16(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP75])])).
fof(f139, plain, ((~ (e21 = h16(e15)) & ~ (e21 = h16(e14)) & ~ (e21 = h16(e13)) & ~ (e21 = h16(e12)) & ~ (e21 = h16(e11)) & ~ (e21 = h16(e10))) | ~ sP76), inference(usedef, [], [e139])).
fof(e139, plain, (sP76 <=> (~ (e21 = h16(e15)) & ~ (e21 = h16(e14)) & ~ (e21 = h16(e13)) & ~ (e21 = h16(e12)) & ~ (e21 = h16(e11)) & ~ (e21 = h16(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP76])])).
fof(f140, plain, ((~ (e22 = h16(e15)) & ~ (e22 = h16(e14)) & ~ (e22 = h16(e13)) & ~ (e22 = h16(e12)) & ~ (e22 = h16(e11)) & ~ (e22 = h16(e10))) | ~ sP77), inference(usedef, [], [e140])).
fof(e140, plain, (sP77 <=> (~ (e22 = h16(e15)) & ~ (e22 = h16(e14)) & ~ (e22 = h16(e13)) & ~ (e22 = h16(e12)) & ~ (e22 = h16(e11)) & ~ (e22 = h16(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP77])])).
fof(f141, plain, ((~ (e23 = h16(e15)) & ~ (e23 = h16(e14)) & ~ (e23 = h16(e13)) & ~ (e23 = h16(e12)) & ~ (e23 = h16(e11)) & ~ (e23 = h16(e10))) | ~ sP78), inference(usedef, [], [e141])).
fof(e141, plain, (sP78 <=> (~ (e23 = h16(e15)) & ~ (e23 = h16(e14)) & ~ (e23 = h16(e13)) & ~ (e23 = h16(e12)) & ~ (e23 = h16(e11)) & ~ (e23 = h16(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP78])])).
fof(f142, plain, ((~ (e24 = h16(e15)) & ~ (e24 = h16(e14)) & ~ (e24 = h16(e13)) & ~ (e24 = h16(e12)) & ~ (e24 = h16(e11)) & ~ (e24 = h16(e10))) | ~ sP79), inference(usedef, [], [e142])).
fof(e142, plain, (sP79 <=> (~ (e24 = h16(e15)) & ~ (e24 = h16(e14)) & ~ (e24 = h16(e13)) & ~ (e24 = h16(e12)) & ~ (e24 = h16(e11)) & ~ (e24 = h16(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP79])])).
fof(f143, plain, ((~ (e20 = h17(e15)) & ~ (e20 = h17(e14)) & ~ (e20 = h17(e13)) & ~ (e20 = h17(e12)) & ~ (e20 = h17(e11)) & ~ (e20 = h17(e10))) | ~ sP80), inference(usedef, [], [e143])).
fof(e143, plain, (sP80 <=> (~ (e20 = h17(e15)) & ~ (e20 = h17(e14)) & ~ (e20 = h17(e13)) & ~ (e20 = h17(e12)) & ~ (e20 = h17(e11)) & ~ (e20 = h17(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP80])])).
fof(f144, plain, ((~ (e21 = h17(e15)) & ~ (e21 = h17(e14)) & ~ (e21 = h17(e13)) & ~ (e21 = h17(e12)) & ~ (e21 = h17(e11)) & ~ (e21 = h17(e10))) | ~ sP81), inference(usedef, [], [e144])).
fof(e144, plain, (sP81 <=> (~ (e21 = h17(e15)) & ~ (e21 = h17(e14)) & ~ (e21 = h17(e13)) & ~ (e21 = h17(e12)) & ~ (e21 = h17(e11)) & ~ (e21 = h17(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP81])])).
fof(f145, plain, ((~ (e22 = h17(e15)) & ~ (e22 = h17(e14)) & ~ (e22 = h17(e13)) & ~ (e22 = h17(e12)) & ~ (e22 = h17(e11)) & ~ (e22 = h17(e10))) | ~ sP82), inference(usedef, [], [e145])).
fof(e145, plain, (sP82 <=> (~ (e22 = h17(e15)) & ~ (e22 = h17(e14)) & ~ (e22 = h17(e13)) & ~ (e22 = h17(e12)) & ~ (e22 = h17(e11)) & ~ (e22 = h17(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP82])])).
fof(f146, plain, ((~ (e23 = h17(e15)) & ~ (e23 = h17(e14)) & ~ (e23 = h17(e13)) & ~ (e23 = h17(e12)) & ~ (e23 = h17(e11)) & ~ (e23 = h17(e10))) | ~ sP83), inference(usedef, [], [e146])).
fof(e146, plain, (sP83 <=> (~ (e23 = h17(e15)) & ~ (e23 = h17(e14)) & ~ (e23 = h17(e13)) & ~ (e23 = h17(e12)) & ~ (e23 = h17(e11)) & ~ (e23 = h17(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP83])])).
fof(f147, plain, ((~ (e24 = h17(e15)) & ~ (e24 = h17(e14)) & ~ (e24 = h17(e13)) & ~ (e24 = h17(e12)) & ~ (e24 = h17(e11)) & ~ (e24 = h17(e10))) | ~ sP84), inference(usedef, [], [e147])).
fof(e147, plain, (sP84 <=> (~ (e24 = h17(e15)) & ~ (e24 = h17(e14)) & ~ (e24 = h17(e13)) & ~ (e24 = h17(e12)) & ~ (e24 = h17(e11)) & ~ (e24 = h17(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP84])])).
fof(f148, plain, ((~ (e20 = h18(e15)) & ~ (e20 = h18(e14)) & ~ (e20 = h18(e13)) & ~ (e20 = h18(e12)) & ~ (e20 = h18(e11)) & ~ (e20 = h18(e10))) | ~ sP85), inference(usedef, [], [e148])).
fof(e148, plain, (sP85 <=> (~ (e20 = h18(e15)) & ~ (e20 = h18(e14)) & ~ (e20 = h18(e13)) & ~ (e20 = h18(e12)) & ~ (e20 = h18(e11)) & ~ (e20 = h18(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP85])])).
fof(f149, plain, ((~ (e21 = h18(e15)) & ~ (e21 = h18(e14)) & ~ (e21 = h18(e13)) & ~ (e21 = h18(e12)) & ~ (e21 = h18(e11)) & ~ (e21 = h18(e10))) | ~ sP86), inference(usedef, [], [e149])).
fof(e149, plain, (sP86 <=> (~ (e21 = h18(e15)) & ~ (e21 = h18(e14)) & ~ (e21 = h18(e13)) & ~ (e21 = h18(e12)) & ~ (e21 = h18(e11)) & ~ (e21 = h18(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP86])])).
fof(f150, plain, ((~ (e22 = h18(e15)) & ~ (e22 = h18(e14)) & ~ (e22 = h18(e13)) & ~ (e22 = h18(e12)) & ~ (e22 = h18(e11)) & ~ (e22 = h18(e10))) | ~ sP87), inference(usedef, [], [e150])).
fof(e150, plain, (sP87 <=> (~ (e22 = h18(e15)) & ~ (e22 = h18(e14)) & ~ (e22 = h18(e13)) & ~ (e22 = h18(e12)) & ~ (e22 = h18(e11)) & ~ (e22 = h18(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP87])])).
fof(f151, plain, ((~ (e23 = h18(e15)) & ~ (e23 = h18(e14)) & ~ (e23 = h18(e13)) & ~ (e23 = h18(e12)) & ~ (e23 = h18(e11)) & ~ (e23 = h18(e10))) | ~ sP88), inference(usedef, [], [e151])).
fof(e151, plain, (sP88 <=> (~ (e23 = h18(e15)) & ~ (e23 = h18(e14)) & ~ (e23 = h18(e13)) & ~ (e23 = h18(e12)) & ~ (e23 = h18(e11)) & ~ (e23 = h18(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP88])])).
fof(f152, plain, ((~ (e24 = h18(e15)) & ~ (e24 = h18(e14)) & ~ (e24 = h18(e13)) & ~ (e24 = h18(e12)) & ~ (e24 = h18(e11)) & ~ (e24 = h18(e10))) | ~ sP89), inference(usedef, [], [e152])).
fof(e152, plain, (sP89 <=> (~ (e24 = h18(e15)) & ~ (e24 = h18(e14)) & ~ (e24 = h18(e13)) & ~ (e24 = h18(e12)) & ~ (e24 = h18(e11)) & ~ (e24 = h18(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP89])])).
fof(f153, plain, ((~ (e20 = h19(e15)) & ~ (e20 = h19(e14)) & ~ (e20 = h19(e13)) & ~ (e20 = h19(e12)) & ~ (e20 = h19(e11)) & ~ (e20 = h19(e10))) | ~ sP90), inference(usedef, [], [e153])).
fof(e153, plain, (sP90 <=> (~ (e20 = h19(e15)) & ~ (e20 = h19(e14)) & ~ (e20 = h19(e13)) & ~ (e20 = h19(e12)) & ~ (e20 = h19(e11)) & ~ (e20 = h19(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP90])])).
fof(f154, plain, ((~ (e21 = h19(e15)) & ~ (e21 = h19(e14)) & ~ (e21 = h19(e13)) & ~ (e21 = h19(e12)) & ~ (e21 = h19(e11)) & ~ (e21 = h19(e10))) | ~ sP91), inference(usedef, [], [e154])).
fof(e154, plain, (sP91 <=> (~ (e21 = h19(e15)) & ~ (e21 = h19(e14)) & ~ (e21 = h19(e13)) & ~ (e21 = h19(e12)) & ~ (e21 = h19(e11)) & ~ (e21 = h19(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP91])])).
fof(f155, plain, ((~ (e22 = h19(e15)) & ~ (e22 = h19(e14)) & ~ (e22 = h19(e13)) & ~ (e22 = h19(e12)) & ~ (e22 = h19(e11)) & ~ (e22 = h19(e10))) | ~ sP92), inference(usedef, [], [e155])).
fof(e155, plain, (sP92 <=> (~ (e22 = h19(e15)) & ~ (e22 = h19(e14)) & ~ (e22 = h19(e13)) & ~ (e22 = h19(e12)) & ~ (e22 = h19(e11)) & ~ (e22 = h19(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP92])])).
fof(f156, plain, ((~ (e23 = h19(e15)) & ~ (e23 = h19(e14)) & ~ (e23 = h19(e13)) & ~ (e23 = h19(e12)) & ~ (e23 = h19(e11)) & ~ (e23 = h19(e10))) | ~ sP93), inference(usedef, [], [e156])).
fof(e156, plain, (sP93 <=> (~ (e23 = h19(e15)) & ~ (e23 = h19(e14)) & ~ (e23 = h19(e13)) & ~ (e23 = h19(e12)) & ~ (e23 = h19(e11)) & ~ (e23 = h19(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP93])])).
fof(f157, plain, ((~ (e24 = h19(e15)) & ~ (e24 = h19(e14)) & ~ (e24 = h19(e13)) & ~ (e24 = h19(e12)) & ~ (e24 = h19(e11)) & ~ (e24 = h19(e10))) | ~ sP94), inference(usedef, [], [e157])).
fof(e157, plain, (sP94 <=> (~ (e24 = h19(e15)) & ~ (e24 = h19(e14)) & ~ (e24 = h19(e13)) & ~ (e24 = h19(e12)) & ~ (e24 = h19(e11)) & ~ (e24 = h19(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP94])])).
fof(f158, plain, ((~ (e20 = h20(e15)) & ~ (e20 = h20(e14)) & ~ (e20 = h20(e13)) & ~ (e20 = h20(e12)) & ~ (e20 = h20(e11)) & ~ (e20 = h20(e10))) | ~ sP95), inference(usedef, [], [e158])).
fof(e158, plain, (sP95 <=> (~ (e20 = h20(e15)) & ~ (e20 = h20(e14)) & ~ (e20 = h20(e13)) & ~ (e20 = h20(e12)) & ~ (e20 = h20(e11)) & ~ (e20 = h20(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP95])])).
fof(f159, plain, ((~ (e21 = h20(e15)) & ~ (e21 = h20(e14)) & ~ (e21 = h20(e13)) & ~ (e21 = h20(e12)) & ~ (e21 = h20(e11)) & ~ (e21 = h20(e10))) | ~ sP96), inference(usedef, [], [e159])).
fof(e159, plain, (sP96 <=> (~ (e21 = h20(e15)) & ~ (e21 = h20(e14)) & ~ (e21 = h20(e13)) & ~ (e21 = h20(e12)) & ~ (e21 = h20(e11)) & ~ (e21 = h20(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP96])])).
fof(f160, plain, ((~ (e22 = h20(e15)) & ~ (e22 = h20(e14)) & ~ (e22 = h20(e13)) & ~ (e22 = h20(e12)) & ~ (e22 = h20(e11)) & ~ (e22 = h20(e10))) | ~ sP97), inference(usedef, [], [e160])).
fof(e160, plain, (sP97 <=> (~ (e22 = h20(e15)) & ~ (e22 = h20(e14)) & ~ (e22 = h20(e13)) & ~ (e22 = h20(e12)) & ~ (e22 = h20(e11)) & ~ (e22 = h20(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP97])])).
fof(f161, plain, ((~ (e23 = h20(e15)) & ~ (e23 = h20(e14)) & ~ (e23 = h20(e13)) & ~ (e23 = h20(e12)) & ~ (e23 = h20(e11)) & ~ (e23 = h20(e10))) | ~ sP98), inference(usedef, [], [e161])).
fof(e161, plain, (sP98 <=> (~ (e23 = h20(e15)) & ~ (e23 = h20(e14)) & ~ (e23 = h20(e13)) & ~ (e23 = h20(e12)) & ~ (e23 = h20(e11)) & ~ (e23 = h20(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP98])])).
fof(f162, plain, ((~ (e24 = h20(e15)) & ~ (e24 = h20(e14)) & ~ (e24 = h20(e13)) & ~ (e24 = h20(e12)) & ~ (e24 = h20(e11)) & ~ (e24 = h20(e10))) | ~ sP99), inference(usedef, [], [e162])).
fof(e162, plain, (sP99 <=> (~ (e24 = h20(e15)) & ~ (e24 = h20(e14)) & ~ (e24 = h20(e13)) & ~ (e24 = h20(e12)) & ~ (e24 = h20(e11)) & ~ (e24 = h20(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP99])])).
fof(f163, plain, ((~ (e20 = h21(e15)) & ~ (e20 = h21(e14)) & ~ (e20 = h21(e13)) & ~ (e20 = h21(e12)) & ~ (e20 = h21(e11)) & ~ (e20 = h21(e10))) | ~ sP100), inference(usedef, [], [e163])).
fof(e163, plain, (sP100 <=> (~ (e20 = h21(e15)) & ~ (e20 = h21(e14)) & ~ (e20 = h21(e13)) & ~ (e20 = h21(e12)) & ~ (e20 = h21(e11)) & ~ (e20 = h21(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP100])])).
fof(f164, plain, ((~ (e21 = h21(e15)) & ~ (e21 = h21(e14)) & ~ (e21 = h21(e13)) & ~ (e21 = h21(e12)) & ~ (e21 = h21(e11)) & ~ (e21 = h21(e10))) | ~ sP101), inference(usedef, [], [e164])).
fof(e164, plain, (sP101 <=> (~ (e21 = h21(e15)) & ~ (e21 = h21(e14)) & ~ (e21 = h21(e13)) & ~ (e21 = h21(e12)) & ~ (e21 = h21(e11)) & ~ (e21 = h21(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP101])])).
fof(f165, plain, ((~ (e22 = h21(e15)) & ~ (e22 = h21(e14)) & ~ (e22 = h21(e13)) & ~ (e22 = h21(e12)) & ~ (e22 = h21(e11)) & ~ (e22 = h21(e10))) | ~ sP102), inference(usedef, [], [e165])).
fof(e165, plain, (sP102 <=> (~ (e22 = h21(e15)) & ~ (e22 = h21(e14)) & ~ (e22 = h21(e13)) & ~ (e22 = h21(e12)) & ~ (e22 = h21(e11)) & ~ (e22 = h21(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP102])])).
fof(f166, plain, ((~ (e23 = h21(e15)) & ~ (e23 = h21(e14)) & ~ (e23 = h21(e13)) & ~ (e23 = h21(e12)) & ~ (e23 = h21(e11)) & ~ (e23 = h21(e10))) | ~ sP103), inference(usedef, [], [e166])).
fof(e166, plain, (sP103 <=> (~ (e23 = h21(e15)) & ~ (e23 = h21(e14)) & ~ (e23 = h21(e13)) & ~ (e23 = h21(e12)) & ~ (e23 = h21(e11)) & ~ (e23 = h21(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP103])])).
fof(f167, plain, ((~ (e24 = h21(e15)) & ~ (e24 = h21(e14)) & ~ (e24 = h21(e13)) & ~ (e24 = h21(e12)) & ~ (e24 = h21(e11)) & ~ (e24 = h21(e10))) | ~ sP104), inference(usedef, [], [e167])).
fof(e167, plain, (sP104 <=> (~ (e24 = h21(e15)) & ~ (e24 = h21(e14)) & ~ (e24 = h21(e13)) & ~ (e24 = h21(e12)) & ~ (e24 = h21(e11)) & ~ (e24 = h21(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP104])])).
fof(f168, plain, ((~ (e20 = h22(e15)) & ~ (e20 = h22(e14)) & ~ (e20 = h22(e13)) & ~ (e20 = h22(e12)) & ~ (e20 = h22(e11)) & ~ (e20 = h22(e10))) | ~ sP105), inference(usedef, [], [e168])).
fof(e168, plain, (sP105 <=> (~ (e20 = h22(e15)) & ~ (e20 = h22(e14)) & ~ (e20 = h22(e13)) & ~ (e20 = h22(e12)) & ~ (e20 = h22(e11)) & ~ (e20 = h22(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP105])])).
fof(f169, plain, ((~ (e21 = h22(e15)) & ~ (e21 = h22(e14)) & ~ (e21 = h22(e13)) & ~ (e21 = h22(e12)) & ~ (e21 = h22(e11)) & ~ (e21 = h22(e10))) | ~ sP106), inference(usedef, [], [e169])).
fof(e169, plain, (sP106 <=> (~ (e21 = h22(e15)) & ~ (e21 = h22(e14)) & ~ (e21 = h22(e13)) & ~ (e21 = h22(e12)) & ~ (e21 = h22(e11)) & ~ (e21 = h22(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP106])])).
fof(f170, plain, ((~ (e22 = h22(e15)) & ~ (e22 = h22(e14)) & ~ (e22 = h22(e13)) & ~ (e22 = h22(e12)) & ~ (e22 = h22(e11)) & ~ (e22 = h22(e10))) | ~ sP107), inference(usedef, [], [e170])).
fof(e170, plain, (sP107 <=> (~ (e22 = h22(e15)) & ~ (e22 = h22(e14)) & ~ (e22 = h22(e13)) & ~ (e22 = h22(e12)) & ~ (e22 = h22(e11)) & ~ (e22 = h22(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP107])])).
fof(f171, plain, ((~ (e23 = h22(e15)) & ~ (e23 = h22(e14)) & ~ (e23 = h22(e13)) & ~ (e23 = h22(e12)) & ~ (e23 = h22(e11)) & ~ (e23 = h22(e10))) | ~ sP108), inference(usedef, [], [e171])).
fof(e171, plain, (sP108 <=> (~ (e23 = h22(e15)) & ~ (e23 = h22(e14)) & ~ (e23 = h22(e13)) & ~ (e23 = h22(e12)) & ~ (e23 = h22(e11)) & ~ (e23 = h22(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP108])])).
fof(f172, plain, ((~ (e24 = h22(e15)) & ~ (e24 = h22(e14)) & ~ (e24 = h22(e13)) & ~ (e24 = h22(e12)) & ~ (e24 = h22(e11)) & ~ (e24 = h22(e10))) | ~ sP109), inference(usedef, [], [e172])).
fof(e172, plain, (sP109 <=> (~ (e24 = h22(e15)) & ~ (e24 = h22(e14)) & ~ (e24 = h22(e13)) & ~ (e24 = h22(e12)) & ~ (e24 = h22(e11)) & ~ (e24 = h22(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP109])])).
fof(f173, plain, ((~ (e20 = h23(e15)) & ~ (e20 = h23(e14)) & ~ (e20 = h23(e13)) & ~ (e20 = h23(e12)) & ~ (e20 = h23(e11)) & ~ (e20 = h23(e10))) | ~ sP110), inference(usedef, [], [e173])).
fof(e173, plain, (sP110 <=> (~ (e20 = h23(e15)) & ~ (e20 = h23(e14)) & ~ (e20 = h23(e13)) & ~ (e20 = h23(e12)) & ~ (e20 = h23(e11)) & ~ (e20 = h23(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP110])])).
fof(f174, plain, ((~ (e21 = h23(e15)) & ~ (e21 = h23(e14)) & ~ (e21 = h23(e13)) & ~ (e21 = h23(e12)) & ~ (e21 = h23(e11)) & ~ (e21 = h23(e10))) | ~ sP111), inference(usedef, [], [e174])).
fof(e174, plain, (sP111 <=> (~ (e21 = h23(e15)) & ~ (e21 = h23(e14)) & ~ (e21 = h23(e13)) & ~ (e21 = h23(e12)) & ~ (e21 = h23(e11)) & ~ (e21 = h23(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP111])])).
fof(f175, plain, ((~ (e22 = h23(e15)) & ~ (e22 = h23(e14)) & ~ (e22 = h23(e13)) & ~ (e22 = h23(e12)) & ~ (e22 = h23(e11)) & ~ (e22 = h23(e10))) | ~ sP112), inference(usedef, [], [e175])).
fof(e175, plain, (sP112 <=> (~ (e22 = h23(e15)) & ~ (e22 = h23(e14)) & ~ (e22 = h23(e13)) & ~ (e22 = h23(e12)) & ~ (e22 = h23(e11)) & ~ (e22 = h23(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP112])])).
fof(f176, plain, ((~ (e23 = h23(e15)) & ~ (e23 = h23(e14)) & ~ (e23 = h23(e13)) & ~ (e23 = h23(e12)) & ~ (e23 = h23(e11)) & ~ (e23 = h23(e10))) | ~ sP113), inference(usedef, [], [e176])).
fof(e176, plain, (sP113 <=> (~ (e23 = h23(e15)) & ~ (e23 = h23(e14)) & ~ (e23 = h23(e13)) & ~ (e23 = h23(e12)) & ~ (e23 = h23(e11)) & ~ (e23 = h23(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP113])])).
fof(f177, plain, ((~ (e24 = h23(e15)) & ~ (e24 = h23(e14)) & ~ (e24 = h23(e13)) & ~ (e24 = h23(e12)) & ~ (e24 = h23(e11)) & ~ (e24 = h23(e10))) | ~ sP114), inference(usedef, [], [e177])).
fof(e177, plain, (sP114 <=> (~ (e24 = h23(e15)) & ~ (e24 = h23(e14)) & ~ (e24 = h23(e13)) & ~ (e24 = h23(e12)) & ~ (e24 = h23(e11)) & ~ (e24 = h23(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP114])])).
fof(f178, plain, ((~ (e20 = h24(e15)) & ~ (e20 = h24(e14)) & ~ (e20 = h24(e13)) & ~ (e20 = h24(e12)) & ~ (e20 = h24(e11)) & ~ (e20 = h24(e10))) | ~ sP115), inference(usedef, [], [e178])).
fof(e178, plain, (sP115 <=> (~ (e20 = h24(e15)) & ~ (e20 = h24(e14)) & ~ (e20 = h24(e13)) & ~ (e20 = h24(e12)) & ~ (e20 = h24(e11)) & ~ (e20 = h24(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP115])])).
fof(f179, plain, ((~ (e21 = h24(e15)) & ~ (e21 = h24(e14)) & ~ (e21 = h24(e13)) & ~ (e21 = h24(e12)) & ~ (e21 = h24(e11)) & ~ (e21 = h24(e10))) | ~ sP116), inference(usedef, [], [e179])).
fof(e179, plain, (sP116 <=> (~ (e21 = h24(e15)) & ~ (e21 = h24(e14)) & ~ (e21 = h24(e13)) & ~ (e21 = h24(e12)) & ~ (e21 = h24(e11)) & ~ (e21 = h24(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP116])])).
fof(f180, plain, ((~ (e22 = h24(e15)) & ~ (e22 = h24(e14)) & ~ (e22 = h24(e13)) & ~ (e22 = h24(e12)) & ~ (e22 = h24(e11)) & ~ (e22 = h24(e10))) | ~ sP117), inference(usedef, [], [e180])).
fof(e180, plain, (sP117 <=> (~ (e22 = h24(e15)) & ~ (e22 = h24(e14)) & ~ (e22 = h24(e13)) & ~ (e22 = h24(e12)) & ~ (e22 = h24(e11)) & ~ (e22 = h24(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP117])])).
fof(f181, plain, ((~ (e23 = h24(e15)) & ~ (e23 = h24(e14)) & ~ (e23 = h24(e13)) & ~ (e23 = h24(e12)) & ~ (e23 = h24(e11)) & ~ (e23 = h24(e10))) | ~ sP118), inference(usedef, [], [e181])).
fof(e181, plain, (sP118 <=> (~ (e23 = h24(e15)) & ~ (e23 = h24(e14)) & ~ (e23 = h24(e13)) & ~ (e23 = h24(e12)) & ~ (e23 = h24(e11)) & ~ (e23 = h24(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP118])])).
fof(f182, plain, ((~ (e24 = h24(e15)) & ~ (e24 = h24(e14)) & ~ (e24 = h24(e13)) & ~ (e24 = h24(e12)) & ~ (e24 = h24(e11)) & ~ (e24 = h24(e10))) | ~ sP119), inference(usedef, [], [e182])).
fof(e182, plain, (sP119 <=> (~ (e24 = h24(e15)) & ~ (e24 = h24(e14)) & ~ (e24 = h24(e13)) & ~ (e24 = h24(e12)) & ~ (e24 = h24(e11)) & ~ (e24 = h24(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP119])])).
fof(f183, plain, ((~ (e20 = h25(e15)) & ~ (e20 = h25(e14)) & ~ (e20 = h25(e13)) & ~ (e20 = h25(e12)) & ~ (e20 = h25(e11)) & ~ (e20 = h25(e10))) | ~ sP120), inference(usedef, [], [e183])).
fof(e183, plain, (sP120 <=> (~ (e20 = h25(e15)) & ~ (e20 = h25(e14)) & ~ (e20 = h25(e13)) & ~ (e20 = h25(e12)) & ~ (e20 = h25(e11)) & ~ (e20 = h25(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP120])])).
fof(f184, plain, ((~ (e21 = h25(e15)) & ~ (e21 = h25(e14)) & ~ (e21 = h25(e13)) & ~ (e21 = h25(e12)) & ~ (e21 = h25(e11)) & ~ (e21 = h25(e10))) | ~ sP121), inference(usedef, [], [e184])).
fof(e184, plain, (sP121 <=> (~ (e21 = h25(e15)) & ~ (e21 = h25(e14)) & ~ (e21 = h25(e13)) & ~ (e21 = h25(e12)) & ~ (e21 = h25(e11)) & ~ (e21 = h25(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP121])])).
fof(f185, plain, ((~ (e22 = h25(e15)) & ~ (e22 = h25(e14)) & ~ (e22 = h25(e13)) & ~ (e22 = h25(e12)) & ~ (e22 = h25(e11)) & ~ (e22 = h25(e10))) | ~ sP122), inference(usedef, [], [e185])).
fof(e185, plain, (sP122 <=> (~ (e22 = h25(e15)) & ~ (e22 = h25(e14)) & ~ (e22 = h25(e13)) & ~ (e22 = h25(e12)) & ~ (e22 = h25(e11)) & ~ (e22 = h25(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP122])])).
fof(f186, plain, ((~ (e23 = h25(e15)) & ~ (e23 = h25(e14)) & ~ (e23 = h25(e13)) & ~ (e23 = h25(e12)) & ~ (e23 = h25(e11)) & ~ (e23 = h25(e10))) | ~ sP123), inference(usedef, [], [e186])).
fof(e186, plain, (sP123 <=> (~ (e23 = h25(e15)) & ~ (e23 = h25(e14)) & ~ (e23 = h25(e13)) & ~ (e23 = h25(e12)) & ~ (e23 = h25(e11)) & ~ (e23 = h25(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP123])])).
fof(f187, plain, ((~ (e24 = h25(e15)) & ~ (e24 = h25(e14)) & ~ (e24 = h25(e13)) & ~ (e24 = h25(e12)) & ~ (e24 = h25(e11)) & ~ (e24 = h25(e10))) | ~ sP124), inference(usedef, [], [e187])).
fof(e187, plain, (sP124 <=> (~ (e24 = h25(e15)) & ~ (e24 = h25(e14)) & ~ (e24 = h25(e13)) & ~ (e24 = h25(e12)) & ~ (e24 = h25(e11)) & ~ (e24 = h25(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP124])])).
fof(f188, plain, ((~ (e20 = h26(e15)) & ~ (e20 = h26(e14)) & ~ (e20 = h26(e13)) & ~ (e20 = h26(e12)) & ~ (e20 = h26(e11)) & ~ (e20 = h26(e10))) | ~ sP125), inference(usedef, [], [e188])).
fof(e188, plain, (sP125 <=> (~ (e20 = h26(e15)) & ~ (e20 = h26(e14)) & ~ (e20 = h26(e13)) & ~ (e20 = h26(e12)) & ~ (e20 = h26(e11)) & ~ (e20 = h26(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP125])])).
fof(f189, plain, ((~ (e21 = h26(e15)) & ~ (e21 = h26(e14)) & ~ (e21 = h26(e13)) & ~ (e21 = h26(e12)) & ~ (e21 = h26(e11)) & ~ (e21 = h26(e10))) | ~ sP126), inference(usedef, [], [e189])).
fof(e189, plain, (sP126 <=> (~ (e21 = h26(e15)) & ~ (e21 = h26(e14)) & ~ (e21 = h26(e13)) & ~ (e21 = h26(e12)) & ~ (e21 = h26(e11)) & ~ (e21 = h26(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP126])])).
fof(f190, plain, ((~ (e22 = h26(e15)) & ~ (e22 = h26(e14)) & ~ (e22 = h26(e13)) & ~ (e22 = h26(e12)) & ~ (e22 = h26(e11)) & ~ (e22 = h26(e10))) | ~ sP127), inference(usedef, [], [e190])).
fof(e190, plain, (sP127 <=> (~ (e22 = h26(e15)) & ~ (e22 = h26(e14)) & ~ (e22 = h26(e13)) & ~ (e22 = h26(e12)) & ~ (e22 = h26(e11)) & ~ (e22 = h26(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP127])])).
fof(f191, plain, ((~ (e23 = h26(e15)) & ~ (e23 = h26(e14)) & ~ (e23 = h26(e13)) & ~ (e23 = h26(e12)) & ~ (e23 = h26(e11)) & ~ (e23 = h26(e10))) | ~ sP128), inference(usedef, [], [e191])).
fof(e191, plain, (sP128 <=> (~ (e23 = h26(e15)) & ~ (e23 = h26(e14)) & ~ (e23 = h26(e13)) & ~ (e23 = h26(e12)) & ~ (e23 = h26(e11)) & ~ (e23 = h26(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP128])])).
fof(f192, plain, ((~ (e24 = h26(e15)) & ~ (e24 = h26(e14)) & ~ (e24 = h26(e13)) & ~ (e24 = h26(e12)) & ~ (e24 = h26(e11)) & ~ (e24 = h26(e10))) | ~ sP129), inference(usedef, [], [e192])).
fof(e192, plain, (sP129 <=> (~ (e24 = h26(e15)) & ~ (e24 = h26(e14)) & ~ (e24 = h26(e13)) & ~ (e24 = h26(e12)) & ~ (e24 = h26(e11)) & ~ (e24 = h26(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP129])])).
fof(f193, plain, ((~ (e20 = h27(e15)) & ~ (e20 = h27(e14)) & ~ (e20 = h27(e13)) & ~ (e20 = h27(e12)) & ~ (e20 = h27(e11)) & ~ (e20 = h27(e10))) | ~ sP130), inference(usedef, [], [e193])).
fof(e193, plain, (sP130 <=> (~ (e20 = h27(e15)) & ~ (e20 = h27(e14)) & ~ (e20 = h27(e13)) & ~ (e20 = h27(e12)) & ~ (e20 = h27(e11)) & ~ (e20 = h27(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP130])])).
fof(f194, plain, ((~ (e21 = h27(e15)) & ~ (e21 = h27(e14)) & ~ (e21 = h27(e13)) & ~ (e21 = h27(e12)) & ~ (e21 = h27(e11)) & ~ (e21 = h27(e10))) | ~ sP131), inference(usedef, [], [e194])).
fof(e194, plain, (sP131 <=> (~ (e21 = h27(e15)) & ~ (e21 = h27(e14)) & ~ (e21 = h27(e13)) & ~ (e21 = h27(e12)) & ~ (e21 = h27(e11)) & ~ (e21 = h27(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP131])])).
fof(f195, plain, ((~ (e22 = h27(e15)) & ~ (e22 = h27(e14)) & ~ (e22 = h27(e13)) & ~ (e22 = h27(e12)) & ~ (e22 = h27(e11)) & ~ (e22 = h27(e10))) | ~ sP132), inference(usedef, [], [e195])).
fof(e195, plain, (sP132 <=> (~ (e22 = h27(e15)) & ~ (e22 = h27(e14)) & ~ (e22 = h27(e13)) & ~ (e22 = h27(e12)) & ~ (e22 = h27(e11)) & ~ (e22 = h27(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP132])])).
fof(f196, plain, ((~ (e23 = h27(e15)) & ~ (e23 = h27(e14)) & ~ (e23 = h27(e13)) & ~ (e23 = h27(e12)) & ~ (e23 = h27(e11)) & ~ (e23 = h27(e10))) | ~ sP133), inference(usedef, [], [e196])).
fof(e196, plain, (sP133 <=> (~ (e23 = h27(e15)) & ~ (e23 = h27(e14)) & ~ (e23 = h27(e13)) & ~ (e23 = h27(e12)) & ~ (e23 = h27(e11)) & ~ (e23 = h27(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP133])])).
fof(f197, plain, ((~ (e24 = h27(e15)) & ~ (e24 = h27(e14)) & ~ (e24 = h27(e13)) & ~ (e24 = h27(e12)) & ~ (e24 = h27(e11)) & ~ (e24 = h27(e10))) | ~ sP134), inference(usedef, [], [e197])).
fof(e197, plain, (sP134 <=> (~ (e24 = h27(e15)) & ~ (e24 = h27(e14)) & ~ (e24 = h27(e13)) & ~ (e24 = h27(e12)) & ~ (e24 = h27(e11)) & ~ (e24 = h27(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP134])])).
fof(f198, plain, ((~ (e20 = h28(e15)) & ~ (e20 = h28(e14)) & ~ (e20 = h28(e13)) & ~ (e20 = h28(e12)) & ~ (e20 = h28(e11)) & ~ (e20 = h28(e10))) | ~ sP135), inference(usedef, [], [e198])).
fof(e198, plain, (sP135 <=> (~ (e20 = h28(e15)) & ~ (e20 = h28(e14)) & ~ (e20 = h28(e13)) & ~ (e20 = h28(e12)) & ~ (e20 = h28(e11)) & ~ (e20 = h28(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP135])])).
fof(f199, plain, ((~ (e21 = h28(e15)) & ~ (e21 = h28(e14)) & ~ (e21 = h28(e13)) & ~ (e21 = h28(e12)) & ~ (e21 = h28(e11)) & ~ (e21 = h28(e10))) | ~ sP136), inference(usedef, [], [e199])).
fof(e199, plain, (sP136 <=> (~ (e21 = h28(e15)) & ~ (e21 = h28(e14)) & ~ (e21 = h28(e13)) & ~ (e21 = h28(e12)) & ~ (e21 = h28(e11)) & ~ (e21 = h28(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP136])])).
fof(f200, plain, ((~ (e22 = h28(e15)) & ~ (e22 = h28(e14)) & ~ (e22 = h28(e13)) & ~ (e22 = h28(e12)) & ~ (e22 = h28(e11)) & ~ (e22 = h28(e10))) | ~ sP137), inference(usedef, [], [e200])).
fof(e200, plain, (sP137 <=> (~ (e22 = h28(e15)) & ~ (e22 = h28(e14)) & ~ (e22 = h28(e13)) & ~ (e22 = h28(e12)) & ~ (e22 = h28(e11)) & ~ (e22 = h28(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP137])])).
fof(f201, plain, ((~ (e23 = h28(e15)) & ~ (e23 = h28(e14)) & ~ (e23 = h28(e13)) & ~ (e23 = h28(e12)) & ~ (e23 = h28(e11)) & ~ (e23 = h28(e10))) | ~ sP138), inference(usedef, [], [e201])).
fof(e201, plain, (sP138 <=> (~ (e23 = h28(e15)) & ~ (e23 = h28(e14)) & ~ (e23 = h28(e13)) & ~ (e23 = h28(e12)) & ~ (e23 = h28(e11)) & ~ (e23 = h28(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP138])])).
fof(f202, plain, ((~ (e24 = h28(e15)) & ~ (e24 = h28(e14)) & ~ (e24 = h28(e13)) & ~ (e24 = h28(e12)) & ~ (e24 = h28(e11)) & ~ (e24 = h28(e10))) | ~ sP139), inference(usedef, [], [e202])).
fof(e202, plain, (sP139 <=> (~ (e24 = h28(e15)) & ~ (e24 = h28(e14)) & ~ (e24 = h28(e13)) & ~ (e24 = h28(e12)) & ~ (e24 = h28(e11)) & ~ (e24 = h28(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP139])])).
fof(f203, plain, ((~ (e20 = h29(e15)) & ~ (e20 = h29(e14)) & ~ (e20 = h29(e13)) & ~ (e20 = h29(e12)) & ~ (e20 = h29(e11)) & ~ (e20 = h29(e10))) | ~ sP140), inference(usedef, [], [e203])).
fof(e203, plain, (sP140 <=> (~ (e20 = h29(e15)) & ~ (e20 = h29(e14)) & ~ (e20 = h29(e13)) & ~ (e20 = h29(e12)) & ~ (e20 = h29(e11)) & ~ (e20 = h29(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP140])])).
fof(f204, plain, ((~ (e21 = h29(e15)) & ~ (e21 = h29(e14)) & ~ (e21 = h29(e13)) & ~ (e21 = h29(e12)) & ~ (e21 = h29(e11)) & ~ (e21 = h29(e10))) | ~ sP141), inference(usedef, [], [e204])).
fof(e204, plain, (sP141 <=> (~ (e21 = h29(e15)) & ~ (e21 = h29(e14)) & ~ (e21 = h29(e13)) & ~ (e21 = h29(e12)) & ~ (e21 = h29(e11)) & ~ (e21 = h29(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP141])])).
fof(f205, plain, ((~ (e22 = h29(e15)) & ~ (e22 = h29(e14)) & ~ (e22 = h29(e13)) & ~ (e22 = h29(e12)) & ~ (e22 = h29(e11)) & ~ (e22 = h29(e10))) | ~ sP142), inference(usedef, [], [e205])).
fof(e205, plain, (sP142 <=> (~ (e22 = h29(e15)) & ~ (e22 = h29(e14)) & ~ (e22 = h29(e13)) & ~ (e22 = h29(e12)) & ~ (e22 = h29(e11)) & ~ (e22 = h29(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP142])])).
fof(f206, plain, ((~ (e23 = h29(e15)) & ~ (e23 = h29(e14)) & ~ (e23 = h29(e13)) & ~ (e23 = h29(e12)) & ~ (e23 = h29(e11)) & ~ (e23 = h29(e10))) | ~ sP143), inference(usedef, [], [e206])).
fof(e206, plain, (sP143 <=> (~ (e23 = h29(e15)) & ~ (e23 = h29(e14)) & ~ (e23 = h29(e13)) & ~ (e23 = h29(e12)) & ~ (e23 = h29(e11)) & ~ (e23 = h29(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP143])])).
fof(f207, plain, ((~ (e24 = h29(e15)) & ~ (e24 = h29(e14)) & ~ (e24 = h29(e13)) & ~ (e24 = h29(e12)) & ~ (e24 = h29(e11)) & ~ (e24 = h29(e10))) | ~ sP144), inference(usedef, [], [e207])).
fof(e207, plain, (sP144 <=> (~ (e24 = h29(e15)) & ~ (e24 = h29(e14)) & ~ (e24 = h29(e13)) & ~ (e24 = h29(e12)) & ~ (e24 = h29(e11)) & ~ (e24 = h29(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP144])])).
fof(f208, plain, ((~ (e20 = h30(e15)) & ~ (e20 = h30(e14)) & ~ (e20 = h30(e13)) & ~ (e20 = h30(e12)) & ~ (e20 = h30(e11)) & ~ (e20 = h30(e10))) | ~ sP145), inference(usedef, [], [e208])).
fof(e208, plain, (sP145 <=> (~ (e20 = h30(e15)) & ~ (e20 = h30(e14)) & ~ (e20 = h30(e13)) & ~ (e20 = h30(e12)) & ~ (e20 = h30(e11)) & ~ (e20 = h30(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP145])])).
fof(f209, plain, ((~ (e21 = h30(e15)) & ~ (e21 = h30(e14)) & ~ (e21 = h30(e13)) & ~ (e21 = h30(e12)) & ~ (e21 = h30(e11)) & ~ (e21 = h30(e10))) | ~ sP146), inference(usedef, [], [e209])).
fof(e209, plain, (sP146 <=> (~ (e21 = h30(e15)) & ~ (e21 = h30(e14)) & ~ (e21 = h30(e13)) & ~ (e21 = h30(e12)) & ~ (e21 = h30(e11)) & ~ (e21 = h30(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP146])])).
fof(f210, plain, ((~ (e22 = h30(e15)) & ~ (e22 = h30(e14)) & ~ (e22 = h30(e13)) & ~ (e22 = h30(e12)) & ~ (e22 = h30(e11)) & ~ (e22 = h30(e10))) | ~ sP147), inference(usedef, [], [e210])).
fof(e210, plain, (sP147 <=> (~ (e22 = h30(e15)) & ~ (e22 = h30(e14)) & ~ (e22 = h30(e13)) & ~ (e22 = h30(e12)) & ~ (e22 = h30(e11)) & ~ (e22 = h30(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP147])])).
fof(f211, plain, ((~ (e23 = h30(e15)) & ~ (e23 = h30(e14)) & ~ (e23 = h30(e13)) & ~ (e23 = h30(e12)) & ~ (e23 = h30(e11)) & ~ (e23 = h30(e10))) | ~ sP148), inference(usedef, [], [e211])).
fof(e211, plain, (sP148 <=> (~ (e23 = h30(e15)) & ~ (e23 = h30(e14)) & ~ (e23 = h30(e13)) & ~ (e23 = h30(e12)) & ~ (e23 = h30(e11)) & ~ (e23 = h30(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP148])])).
fof(f212, plain, ((~ (e24 = h30(e15)) & ~ (e24 = h30(e14)) & ~ (e24 = h30(e13)) & ~ (e24 = h30(e12)) & ~ (e24 = h30(e11)) & ~ (e24 = h30(e10))) | ~ sP149), inference(usedef, [], [e212])).
fof(e212, plain, (sP149 <=> (~ (e24 = h30(e15)) & ~ (e24 = h30(e14)) & ~ (e24 = h30(e13)) & ~ (e24 = h30(e12)) & ~ (e24 = h30(e11)) & ~ (e24 = h30(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP149])])).
fof(f62, plain, (((~ (e25 = h30(e15)) & ~ (e25 = h30(e14)) & ~ (e25 = h30(e13)) & ~ (e25 = h30(e12)) & ~ (e25 = h30(e11)) & ~ (e25 = h30(e10))) | (~ (e24 = h30(e15)) & ~ (e24 = h30(e14)) & ~ (e24 = h30(e13)) & ~ (e24 = h30(e12)) & ~ (e24 = h30(e11)) & ~ (e24 = h30(e10))) | (~ (e23 = h30(e15)) & ~ (e23 = h30(e14)) & ~ (e23 = h30(e13)) & ~ (e23 = h30(e12)) & ~ (e23 = h30(e11)) & ~ (e23 = h30(e10))) | (~ (e22 = h30(e15)) & ~ (e22 = h30(e14)) & ~ (e22 = h30(e13)) & ~ (e22 = h30(e12)) & ~ (e22 = h30(e11)) & ~ (e22 = h30(e10))) | (~ (e21 = h30(e15)) & ~ (e21 = h30(e14)) & ~ (e21 = h30(e13)) & ~ (e21 = h30(e12)) & ~ (e21 = h30(e11)) & ~ (e21 = h30(e10))) | (~ (e20 = h30(e15)) & ~ (e20 = h30(e14)) & ~ (e20 = h30(e13)) & ~ (e20 = h30(e12)) & ~ (e20 = h30(e11)) & ~ (e20 = h30(e10))) | ~ (h30(op1(e15, e15)) = op2(h30(e15), h30(e15))) | ~ (h30(op1(e15, e14)) = op2(h30(e15), h30(e14))) | ~ (h30(op1(e15, e13)) = op2(h30(e15), h30(e13))) | ~ (h30(op1(e15, e12)) = op2(h30(e15), h30(e12))) | ~ (h30(op1(e15, e11)) = op2(h30(e15), h30(e11))) | ~ (h30(op1(e15, e10)) = op2(h30(e15), h30(e10))) | ~ (h30(op1(e14, e15)) = op2(h30(e14), h30(e15))) | ~ (h30(op1(e14, e14)) = op2(h30(e14), h30(e14))) | ~ (h30(op1(e14, e13)) = op2(h30(e14), h30(e13))) | ~ (h30(op1(e14, e12)) = op2(h30(e14), h30(e12))) | ~ (h30(op1(e14, e11)) = op2(h30(e14), h30(e11))) | ~ (h30(op1(e14, e10)) = op2(h30(e14), h30(e10))) | ~ (h30(op1(e13, e15)) = op2(h30(e13), h30(e15))) | ~ (h30(op1(e13, e14)) = op2(h30(e13), h30(e14))) | ~ (h30(op1(e13, e13)) = op2(h30(e13), h30(e13))) | ~ (h30(op1(e13, e12)) = op2(h30(e13), h30(e12))) | ~ (h30(op1(e13, e11)) = op2(h30(e13), h30(e11))) | ~ (h30(op1(e13, e10)) = op2(h30(e13), h30(e10))) | ~ (h30(op1(e12, e15)) = op2(h30(e12), h30(e15))) | ~ (h30(op1(e12, e14)) = op2(h30(e12), h30(e14))) | ~ (h30(op1(e12, e13)) = op2(h30(e12), h30(e13))) | ~ (h30(op1(e12, e12)) = op2(h30(e12), h30(e12))) | ~ (h30(op1(e12, e11)) = op2(h30(e12), h30(e11))) | ~ (h30(op1(e12, e10)) = op2(h30(e12), h30(e10))) | ~ (h30(op1(e11, e15)) = op2(h30(e11), h30(e15))) | ~ (h30(op1(e11, e14)) = op2(h30(e11), h30(e14))) | ~ (h30(op1(e11, e13)) = op2(h30(e11), h30(e13))) | ~ (h30(op1(e11, e12)) = op2(h30(e11), h30(e12))) | ~ (h30(op1(e11, e11)) = op2(h30(e11), h30(e11))) | ~ (h30(op1(e11, e10)) = op2(h30(e11), h30(e10))) | ~ (h30(op1(e10, e15)) = op2(h30(e10), h30(e15))) | ~ (h30(op1(e10, e14)) = op2(h30(e10), h30(e14))) | ~ (h30(op1(e10, e13)) = op2(h30(e10), h30(e13))) | ~ (h30(op1(e10, e12)) = op2(h30(e10), h30(e12))) | ~ (h30(op1(e10, e11)) = op2(h30(e10), h30(e11))) | ~ (h30(op1(e10, e10)) = op2(h30(e10), h30(e10)))) & ((~ (e25 = h29(e15)) & ~ (e25 = h29(e14)) & ~ (e25 = h29(e13)) & ~ (e25 = h29(e12)) & ~ (e25 = h29(e11)) & ~ (e25 = h29(e10))) | (~ (e24 = h29(e15)) & ~ (e24 = h29(e14)) & ~ (e24 = h29(e13)) & ~ (e24 = h29(e12)) & ~ (e24 = h29(e11)) & ~ (e24 = h29(e10))) | (~ (e23 = h29(e15)) & ~ (e23 = h29(e14)) & ~ (e23 = h29(e13)) & ~ (e23 = h29(e12)) & ~ (e23 = h29(e11)) & ~ (e23 = h29(e10))) | (~ (e22 = h29(e15)) & ~ (e22 = h29(e14)) & ~ (e22 = h29(e13)) & ~ (e22 = h29(e12)) & ~ (e22 = h29(e11)) & ~ (e22 = h29(e10))) | (~ (e21 = h29(e15)) & ~ (e21 = h29(e14)) & ~ (e21 = h29(e13)) & ~ (e21 = h29(e12)) & ~ (e21 = h29(e11)) & ~ (e21 = h29(e10))) | (~ (e20 = h29(e15)) & ~ (e20 = h29(e14)) & ~ (e20 = h29(e13)) & ~ (e20 = h29(e12)) & ~ (e20 = h29(e11)) & ~ (e20 = h29(e10))) | ~ (h29(op1(e15, e15)) = op2(h29(e15), h29(e15))) | ~ (h29(op1(e15, e14)) = op2(h29(e15), h29(e14))) | ~ (h29(op1(e15, e13)) = op2(h29(e15), h29(e13))) | ~ (h29(op1(e15, e12)) = op2(h29(e15), h29(e12))) | ~ (h29(op1(e15, e11)) = op2(h29(e15), h29(e11))) | ~ (h29(op1(e15, e10)) = op2(h29(e15), h29(e10))) | ~ (h29(op1(e14, e15)) = op2(h29(e14), h29(e15))) | ~ (h29(op1(e14, e14)) = op2(h29(e14), h29(e14))) | ~ (h29(op1(e14, e13)) = op2(h29(e14), h29(e13))) | ~ (h29(op1(e14, e12)) = op2(h29(e14), h29(e12))) | ~ (h29(op1(e14, e11)) = op2(h29(e14), h29(e11))) | ~ (h29(op1(e14, e10)) = op2(h29(e14), h29(e10))) | ~ (h29(op1(e13, e15)) = op2(h29(e13), h29(e15))) | ~ (h29(op1(e13, e14)) = op2(h29(e13), h29(e14))) | ~ (h29(op1(e13, e13)) = op2(h29(e13), h29(e13))) | ~ (h29(op1(e13, e12)) = op2(h29(e13), h29(e12))) | ~ (h29(op1(e13, e11)) = op2(h29(e13), h29(e11))) | ~ (h29(op1(e13, e10)) = op2(h29(e13), h29(e10))) | ~ (h29(op1(e12, e15)) = op2(h29(e12), h29(e15))) | ~ (h29(op1(e12, e14)) = op2(h29(e12), h29(e14))) | ~ (h29(op1(e12, e13)) = op2(h29(e12), h29(e13))) | ~ (h29(op1(e12, e12)) = op2(h29(e12), h29(e12))) | ~ (h29(op1(e12, e11)) = op2(h29(e12), h29(e11))) | ~ (h29(op1(e12, e10)) = op2(h29(e12), h29(e10))) | ~ (h29(op1(e11, e15)) = op2(h29(e11), h29(e15))) | ~ (h29(op1(e11, e14)) = op2(h29(e11), h29(e14))) | ~ (h29(op1(e11, e13)) = op2(h29(e11), h29(e13))) | ~ (h29(op1(e11, e12)) = op2(h29(e11), h29(e12))) | ~ (h29(op1(e11, e11)) = op2(h29(e11), h29(e11))) | ~ (h29(op1(e11, e10)) = op2(h29(e11), h29(e10))) | ~ (h29(op1(e10, e15)) = op2(h29(e10), h29(e15))) | ~ (h29(op1(e10, e14)) = op2(h29(e10), h29(e14))) | ~ (h29(op1(e10, e13)) = op2(h29(e10), h29(e13))) | ~ (h29(op1(e10, e12)) = op2(h29(e10), h29(e12))) | ~ (h29(op1(e10, e11)) = op2(h29(e10), h29(e11))) | ~ (h29(op1(e10, e10)) = op2(h29(e10), h29(e10)))) & ((~ (e25 = h28(e15)) & ~ (e25 = h28(e14)) & ~ (e25 = h28(e13)) & ~ (e25 = h28(e12)) & ~ (e25 = h28(e11)) & ~ (e25 = h28(e10))) | (~ (e24 = h28(e15)) & ~ (e24 = h28(e14)) & ~ (e24 = h28(e13)) & ~ (e24 = h28(e12)) & ~ (e24 = h28(e11)) & ~ (e24 = h28(e10))) | (~ (e23 = h28(e15)) & ~ (e23 = h28(e14)) & ~ (e23 = h28(e13)) & ~ (e23 = h28(e12)) & ~ (e23 = h28(e11)) & ~ (e23 = h28(e10))) | (~ (e22 = h28(e15)) & ~ (e22 = h28(e14)) & ~ (e22 = h28(e13)) & ~ (e22 = h28(e12)) & ~ (e22 = h28(e11)) & ~ (e22 = h28(e10))) | (~ (e21 = h28(e15)) & ~ (e21 = h28(e14)) & ~ (e21 = h28(e13)) & ~ (e21 = h28(e12)) & ~ (e21 = h28(e11)) & ~ (e21 = h28(e10))) | (~ (e20 = h28(e15)) & ~ (e20 = h28(e14)) & ~ (e20 = h28(e13)) & ~ (e20 = h28(e12)) & ~ (e20 = h28(e11)) & ~ (e20 = h28(e10))) | ~ (h28(op1(e15, e15)) = op2(h28(e15), h28(e15))) | ~ (h28(op1(e15, e14)) = op2(h28(e15), h28(e14))) | ~ (h28(op1(e15, e13)) = op2(h28(e15), h28(e13))) | ~ (h28(op1(e15, e12)) = op2(h28(e15), h28(e12))) | ~ (h28(op1(e15, e11)) = op2(h28(e15), h28(e11))) | ~ (h28(op1(e15, e10)) = op2(h28(e15), h28(e10))) | ~ (h28(op1(e14, e15)) = op2(h28(e14), h28(e15))) | ~ (h28(op1(e14, e14)) = op2(h28(e14), h28(e14))) | ~ (h28(op1(e14, e13)) = op2(h28(e14), h28(e13))) | ~ (h28(op1(e14, e12)) = op2(h28(e14), h28(e12))) | ~ (h28(op1(e14, e11)) = op2(h28(e14), h28(e11))) | ~ (h28(op1(e14, e10)) = op2(h28(e14), h28(e10))) | ~ (h28(op1(e13, e15)) = op2(h28(e13), h28(e15))) | ~ (h28(op1(e13, e14)) = op2(h28(e13), h28(e14))) | ~ (h28(op1(e13, e13)) = op2(h28(e13), h28(e13))) | ~ (h28(op1(e13, e12)) = op2(h28(e13), h28(e12))) | ~ (h28(op1(e13, e11)) = op2(h28(e13), h28(e11))) | ~ (h28(op1(e13, e10)) = op2(h28(e13), h28(e10))) | ~ (h28(op1(e12, e15)) = op2(h28(e12), h28(e15))) | ~ (h28(op1(e12, e14)) = op2(h28(e12), h28(e14))) | ~ (h28(op1(e12, e13)) = op2(h28(e12), h28(e13))) | ~ (h28(op1(e12, e12)) = op2(h28(e12), h28(e12))) | ~ (h28(op1(e12, e11)) = op2(h28(e12), h28(e11))) | ~ (h28(op1(e12, e10)) = op2(h28(e12), h28(e10))) | ~ (h28(op1(e11, e15)) = op2(h28(e11), h28(e15))) | ~ (h28(op1(e11, e14)) = op2(h28(e11), h28(e14))) | ~ (h28(op1(e11, e13)) = op2(h28(e11), h28(e13))) | ~ (h28(op1(e11, e12)) = op2(h28(e11), h28(e12))) | ~ (h28(op1(e11, e11)) = op2(h28(e11), h28(e11))) | ~ (h28(op1(e11, e10)) = op2(h28(e11), h28(e10))) | ~ (h28(op1(e10, e15)) = op2(h28(e10), h28(e15))) | ~ (h28(op1(e10, e14)) = op2(h28(e10), h28(e14))) | ~ (h28(op1(e10, e13)) = op2(h28(e10), h28(e13))) | ~ (h28(op1(e10, e12)) = op2(h28(e10), h28(e12))) | ~ (h28(op1(e10, e11)) = op2(h28(e10), h28(e11))) | ~ (h28(op1(e10, e10)) = op2(h28(e10), h28(e10)))) & ((~ (e25 = h27(e15)) & ~ (e25 = h27(e14)) & ~ (e25 = h27(e13)) & ~ (e25 = h27(e12)) & ~ (e25 = h27(e11)) & ~ (e25 = h27(e10))) | (~ (e24 = h27(e15)) & ~ (e24 = h27(e14)) & ~ (e24 = h27(e13)) & ~ (e24 = h27(e12)) & ~ (e24 = h27(e11)) & ~ (e24 = h27(e10))) | (~ (e23 = h27(e15)) & ~ (e23 = h27(e14)) & ~ (e23 = h27(e13)) & ~ (e23 = h27(e12)) & ~ (e23 = h27(e11)) & ~ (e23 = h27(e10))) | (~ (e22 = h27(e15)) & ~ (e22 = h27(e14)) & ~ (e22 = h27(e13)) & ~ (e22 = h27(e12)) & ~ (e22 = h27(e11)) & ~ (e22 = h27(e10))) | (~ (e21 = h27(e15)) & ~ (e21 = h27(e14)) & ~ (e21 = h27(e13)) & ~ (e21 = h27(e12)) & ~ (e21 = h27(e11)) & ~ (e21 = h27(e10))) | (~ (e20 = h27(e15)) & ~ (e20 = h27(e14)) & ~ (e20 = h27(e13)) & ~ (e20 = h27(e12)) & ~ (e20 = h27(e11)) & ~ (e20 = h27(e10))) | ~ (h27(op1(e15, e15)) = op2(h27(e15), h27(e15))) | ~ (h27(op1(e15, e14)) = op2(h27(e15), h27(e14))) | ~ (h27(op1(e15, e13)) = op2(h27(e15), h27(e13))) | ~ (h27(op1(e15, e12)) = op2(h27(e15), h27(e12))) | ~ (h27(op1(e15, e11)) = op2(h27(e15), h27(e11))) | ~ (h27(op1(e15, e10)) = op2(h27(e15), h27(e10))) | ~ (h27(op1(e14, e15)) = op2(h27(e14), h27(e15))) | ~ (h27(op1(e14, e14)) = op2(h27(e14), h27(e14))) | ~ (h27(op1(e14, e13)) = op2(h27(e14), h27(e13))) | ~ (h27(op1(e14, e12)) = op2(h27(e14), h27(e12))) | ~ (h27(op1(e14, e11)) = op2(h27(e14), h27(e11))) | ~ (h27(op1(e14, e10)) = op2(h27(e14), h27(e10))) | ~ (h27(op1(e13, e15)) = op2(h27(e13), h27(e15))) | ~ (h27(op1(e13, e14)) = op2(h27(e13), h27(e14))) | ~ (h27(op1(e13, e13)) = op2(h27(e13), h27(e13))) | ~ (h27(op1(e13, e12)) = op2(h27(e13), h27(e12))) | ~ (h27(op1(e13, e11)) = op2(h27(e13), h27(e11))) | ~ (h27(op1(e13, e10)) = op2(h27(e13), h27(e10))) | ~ (h27(op1(e12, e15)) = op2(h27(e12), h27(e15))) | ~ (h27(op1(e12, e14)) = op2(h27(e12), h27(e14))) | ~ (h27(op1(e12, e13)) = op2(h27(e12), h27(e13))) | ~ (h27(op1(e12, e12)) = op2(h27(e12), h27(e12))) | ~ (h27(op1(e12, e11)) = op2(h27(e12), h27(e11))) | ~ (h27(op1(e12, e10)) = op2(h27(e12), h27(e10))) | ~ (h27(op1(e11, e15)) = op2(h27(e11), h27(e15))) | ~ (h27(op1(e11, e14)) = op2(h27(e11), h27(e14))) | ~ (h27(op1(e11, e13)) = op2(h27(e11), h27(e13))) | ~ (h27(op1(e11, e12)) = op2(h27(e11), h27(e12))) | ~ (h27(op1(e11, e11)) = op2(h27(e11), h27(e11))) | ~ (h27(op1(e11, e10)) = op2(h27(e11), h27(e10))) | ~ (h27(op1(e10, e15)) = op2(h27(e10), h27(e15))) | ~ (h27(op1(e10, e14)) = op2(h27(e10), h27(e14))) | ~ (h27(op1(e10, e13)) = op2(h27(e10), h27(e13))) | ~ (h27(op1(e10, e12)) = op2(h27(e10), h27(e12))) | ~ (h27(op1(e10, e11)) = op2(h27(e10), h27(e11))) | ~ (h27(op1(e10, e10)) = op2(h27(e10), h27(e10)))) & ((~ (e25 = h26(e15)) & ~ (e25 = h26(e14)) & ~ (e25 = h26(e13)) & ~ (e25 = h26(e12)) & ~ (e25 = h26(e11)) & ~ (e25 = h26(e10))) | (~ (e24 = h26(e15)) & ~ (e24 = h26(e14)) & ~ (e24 = h26(e13)) & ~ (e24 = h26(e12)) & ~ (e24 = h26(e11)) & ~ (e24 = h26(e10))) | (~ (e23 = h26(e15)) & ~ (e23 = h26(e14)) & ~ (e23 = h26(e13)) & ~ (e23 = h26(e12)) & ~ (e23 = h26(e11)) & ~ (e23 = h26(e10))) | (~ (e22 = h26(e15)) & ~ (e22 = h26(e14)) & ~ (e22 = h26(e13)) & ~ (e22 = h26(e12)) & ~ (e22 = h26(e11)) & ~ (e22 = h26(e10))) | (~ (e21 = h26(e15)) & ~ (e21 = h26(e14)) & ~ (e21 = h26(e13)) & ~ (e21 = h26(e12)) & ~ (e21 = h26(e11)) & ~ (e21 = h26(e10))) | (~ (e20 = h26(e15)) & ~ (e20 = h26(e14)) & ~ (e20 = h26(e13)) & ~ (e20 = h26(e12)) & ~ (e20 = h26(e11)) & ~ (e20 = h26(e10))) | ~ (h26(op1(e15, e15)) = op2(h26(e15), h26(e15))) | ~ (h26(op1(e15, e14)) = op2(h26(e15), h26(e14))) | ~ (h26(op1(e15, e13)) = op2(h26(e15), h26(e13))) | ~ (h26(op1(e15, e12)) = op2(h26(e15), h26(e12))) | ~ (h26(op1(e15, e11)) = op2(h26(e15), h26(e11))) | ~ (h26(op1(e15, e10)) = op2(h26(e15), h26(e10))) | ~ (h26(op1(e14, e15)) = op2(h26(e14), h26(e15))) | ~ (h26(op1(e14, e14)) = op2(h26(e14), h26(e14))) | ~ (h26(op1(e14, e13)) = op2(h26(e14), h26(e13))) | ~ (h26(op1(e14, e12)) = op2(h26(e14), h26(e12))) | ~ (h26(op1(e14, e11)) = op2(h26(e14), h26(e11))) | ~ (h26(op1(e14, e10)) = op2(h26(e14), h26(e10))) | ~ (h26(op1(e13, e15)) = op2(h26(e13), h26(e15))) | ~ (h26(op1(e13, e14)) = op2(h26(e13), h26(e14))) | ~ (h26(op1(e13, e13)) = op2(h26(e13), h26(e13))) | ~ (h26(op1(e13, e12)) = op2(h26(e13), h26(e12))) | ~ (h26(op1(e13, e11)) = op2(h26(e13), h26(e11))) | ~ (h26(op1(e13, e10)) = op2(h26(e13), h26(e10))) | ~ (h26(op1(e12, e15)) = op2(h26(e12), h26(e15))) | ~ (h26(op1(e12, e14)) = op2(h26(e12), h26(e14))) | ~ (h26(op1(e12, e13)) = op2(h26(e12), h26(e13))) | ~ (h26(op1(e12, e12)) = op2(h26(e12), h26(e12))) | ~ (h26(op1(e12, e11)) = op2(h26(e12), h26(e11))) | ~ (h26(op1(e12, e10)) = op2(h26(e12), h26(e10))) | ~ (h26(op1(e11, e15)) = op2(h26(e11), h26(e15))) | ~ (h26(op1(e11, e14)) = op2(h26(e11), h26(e14))) | ~ (h26(op1(e11, e13)) = op2(h26(e11), h26(e13))) | ~ (h26(op1(e11, e12)) = op2(h26(e11), h26(e12))) | ~ (h26(op1(e11, e11)) = op2(h26(e11), h26(e11))) | ~ (h26(op1(e11, e10)) = op2(h26(e11), h26(e10))) | ~ (h26(op1(e10, e15)) = op2(h26(e10), h26(e15))) | ~ (h26(op1(e10, e14)) = op2(h26(e10), h26(e14))) | ~ (h26(op1(e10, e13)) = op2(h26(e10), h26(e13))) | ~ (h26(op1(e10, e12)) = op2(h26(e10), h26(e12))) | ~ (h26(op1(e10, e11)) = op2(h26(e10), h26(e11))) | ~ (h26(op1(e10, e10)) = op2(h26(e10), h26(e10)))) & ((~ (e25 = h25(e15)) & ~ (e25 = h25(e14)) & ~ (e25 = h25(e13)) & ~ (e25 = h25(e12)) & ~ (e25 = h25(e11)) & ~ (e25 = h25(e10))) | (~ (e24 = h25(e15)) & ~ (e24 = h25(e14)) & ~ (e24 = h25(e13)) & ~ (e24 = h25(e12)) & ~ (e24 = h25(e11)) & ~ (e24 = h25(e10))) | (~ (e23 = h25(e15)) & ~ (e23 = h25(e14)) & ~ (e23 = h25(e13)) & ~ (e23 = h25(e12)) & ~ (e23 = h25(e11)) & ~ (e23 = h25(e10))) | (~ (e22 = h25(e15)) & ~ (e22 = h25(e14)) & ~ (e22 = h25(e13)) & ~ (e22 = h25(e12)) & ~ (e22 = h25(e11)) & ~ (e22 = h25(e10))) | (~ (e21 = h25(e15)) & ~ (e21 = h25(e14)) & ~ (e21 = h25(e13)) & ~ (e21 = h25(e12)) & ~ (e21 = h25(e11)) & ~ (e21 = h25(e10))) | (~ (e20 = h25(e15)) & ~ (e20 = h25(e14)) & ~ (e20 = h25(e13)) & ~ (e20 = h25(e12)) & ~ (e20 = h25(e11)) & ~ (e20 = h25(e10))) | ~ (h25(op1(e15, e15)) = op2(h25(e15), h25(e15))) | ~ (h25(op1(e15, e14)) = op2(h25(e15), h25(e14))) | ~ (h25(op1(e15, e13)) = op2(h25(e15), h25(e13))) | ~ (h25(op1(e15, e12)) = op2(h25(e15), h25(e12))) | ~ (h25(op1(e15, e11)) = op2(h25(e15), h25(e11))) | ~ (h25(op1(e15, e10)) = op2(h25(e15), h25(e10))) | ~ (h25(op1(e14, e15)) = op2(h25(e14), h25(e15))) | ~ (h25(op1(e14, e14)) = op2(h25(e14), h25(e14))) | ~ (h25(op1(e14, e13)) = op2(h25(e14), h25(e13))) | ~ (h25(op1(e14, e12)) = op2(h25(e14), h25(e12))) | ~ (h25(op1(e14, e11)) = op2(h25(e14), h25(e11))) | ~ (h25(op1(e14, e10)) = op2(h25(e14), h25(e10))) | ~ (h25(op1(e13, e15)) = op2(h25(e13), h25(e15))) | ~ (h25(op1(e13, e14)) = op2(h25(e13), h25(e14))) | ~ (h25(op1(e13, e13)) = op2(h25(e13), h25(e13))) | ~ (h25(op1(e13, e12)) = op2(h25(e13), h25(e12))) | ~ (h25(op1(e13, e11)) = op2(h25(e13), h25(e11))) | ~ (h25(op1(e13, e10)) = op2(h25(e13), h25(e10))) | ~ (h25(op1(e12, e15)) = op2(h25(e12), h25(e15))) | ~ (h25(op1(e12, e14)) = op2(h25(e12), h25(e14))) | ~ (h25(op1(e12, e13)) = op2(h25(e12), h25(e13))) | ~ (h25(op1(e12, e12)) = op2(h25(e12), h25(e12))) | ~ (h25(op1(e12, e11)) = op2(h25(e12), h25(e11))) | ~ (h25(op1(e12, e10)) = op2(h25(e12), h25(e10))) | ~ (h25(op1(e11, e15)) = op2(h25(e11), h25(e15))) | ~ (h25(op1(e11, e14)) = op2(h25(e11), h25(e14))) | ~ (h25(op1(e11, e13)) = op2(h25(e11), h25(e13))) | ~ (h25(op1(e11, e12)) = op2(h25(e11), h25(e12))) | ~ (h25(op1(e11, e11)) = op2(h25(e11), h25(e11))) | ~ (h25(op1(e11, e10)) = op2(h25(e11), h25(e10))) | ~ (h25(op1(e10, e15)) = op2(h25(e10), h25(e15))) | ~ (h25(op1(e10, e14)) = op2(h25(e10), h25(e14))) | ~ (h25(op1(e10, e13)) = op2(h25(e10), h25(e13))) | ~ (h25(op1(e10, e12)) = op2(h25(e10), h25(e12))) | ~ (h25(op1(e10, e11)) = op2(h25(e10), h25(e11))) | ~ (h25(op1(e10, e10)) = op2(h25(e10), h25(e10)))) & ((~ (e25 = h24(e15)) & ~ (e25 = h24(e14)) & ~ (e25 = h24(e13)) & ~ (e25 = h24(e12)) & ~ (e25 = h24(e11)) & ~ (e25 = h24(e10))) | (~ (e24 = h24(e15)) & ~ (e24 = h24(e14)) & ~ (e24 = h24(e13)) & ~ (e24 = h24(e12)) & ~ (e24 = h24(e11)) & ~ (e24 = h24(e10))) | (~ (e23 = h24(e15)) & ~ (e23 = h24(e14)) & ~ (e23 = h24(e13)) & ~ (e23 = h24(e12)) & ~ (e23 = h24(e11)) & ~ (e23 = h24(e10))) | (~ (e22 = h24(e15)) & ~ (e22 = h24(e14)) & ~ (e22 = h24(e13)) & ~ (e22 = h24(e12)) & ~ (e22 = h24(e11)) & ~ (e22 = h24(e10))) | (~ (e21 = h24(e15)) & ~ (e21 = h24(e14)) & ~ (e21 = h24(e13)) & ~ (e21 = h24(e12)) & ~ (e21 = h24(e11)) & ~ (e21 = h24(e10))) | (~ (e20 = h24(e15)) & ~ (e20 = h24(e14)) & ~ (e20 = h24(e13)) & ~ (e20 = h24(e12)) & ~ (e20 = h24(e11)) & ~ (e20 = h24(e10))) | ~ (h24(op1(e15, e15)) = op2(h24(e15), h24(e15))) | ~ (h24(op1(e15, e14)) = op2(h24(e15), h24(e14))) | ~ (h24(op1(e15, e13)) = op2(h24(e15), h24(e13))) | ~ (h24(op1(e15, e12)) = op2(h24(e15), h24(e12))) | ~ (h24(op1(e15, e11)) = op2(h24(e15), h24(e11))) | ~ (h24(op1(e15, e10)) = op2(h24(e15), h24(e10))) | ~ (h24(op1(e14, e15)) = op2(h24(e14), h24(e15))) | ~ (h24(op1(e14, e14)) = op2(h24(e14), h24(e14))) | ~ (h24(op1(e14, e13)) = op2(h24(e14), h24(e13))) | ~ (h24(op1(e14, e12)) = op2(h24(e14), h24(e12))) | ~ (h24(op1(e14, e11)) = op2(h24(e14), h24(e11))) | ~ (h24(op1(e14, e10)) = op2(h24(e14), h24(e10))) | ~ (h24(op1(e13, e15)) = op2(h24(e13), h24(e15))) | ~ (h24(op1(e13, e14)) = op2(h24(e13), h24(e14))) | ~ (h24(op1(e13, e13)) = op2(h24(e13), h24(e13))) | ~ (h24(op1(e13, e12)) = op2(h24(e13), h24(e12))) | ~ (h24(op1(e13, e11)) = op2(h24(e13), h24(e11))) | ~ (h24(op1(e13, e10)) = op2(h24(e13), h24(e10))) | ~ (h24(op1(e12, e15)) = op2(h24(e12), h24(e15))) | ~ (h24(op1(e12, e14)) = op2(h24(e12), h24(e14))) | ~ (h24(op1(e12, e13)) = op2(h24(e12), h24(e13))) | ~ (h24(op1(e12, e12)) = op2(h24(e12), h24(e12))) | ~ (h24(op1(e12, e11)) = op2(h24(e12), h24(e11))) | ~ (h24(op1(e12, e10)) = op2(h24(e12), h24(e10))) | ~ (h24(op1(e11, e15)) = op2(h24(e11), h24(e15))) | ~ (h24(op1(e11, e14)) = op2(h24(e11), h24(e14))) | ~ (h24(op1(e11, e13)) = op2(h24(e11), h24(e13))) | ~ (h24(op1(e11, e12)) = op2(h24(e11), h24(e12))) | ~ (h24(op1(e11, e11)) = op2(h24(e11), h24(e11))) | ~ (h24(op1(e11, e10)) = op2(h24(e11), h24(e10))) | ~ (h24(op1(e10, e15)) = op2(h24(e10), h24(e15))) | ~ (h24(op1(e10, e14)) = op2(h24(e10), h24(e14))) | ~ (h24(op1(e10, e13)) = op2(h24(e10), h24(e13))) | ~ (h24(op1(e10, e12)) = op2(h24(e10), h24(e12))) | ~ (h24(op1(e10, e11)) = op2(h24(e10), h24(e11))) | ~ (h24(op1(e10, e10)) = op2(h24(e10), h24(e10)))) & ((~ (e25 = h23(e15)) & ~ (e25 = h23(e14)) & ~ (e25 = h23(e13)) & ~ (e25 = h23(e12)) & ~ (e25 = h23(e11)) & ~ (e25 = h23(e10))) | (~ (e24 = h23(e15)) & ~ (e24 = h23(e14)) & ~ (e24 = h23(e13)) & ~ (e24 = h23(e12)) & ~ (e24 = h23(e11)) & ~ (e24 = h23(e10))) | (~ (e23 = h23(e15)) & ~ (e23 = h23(e14)) & ~ (e23 = h23(e13)) & ~ (e23 = h23(e12)) & ~ (e23 = h23(e11)) & ~ (e23 = h23(e10))) | (~ (e22 = h23(e15)) & ~ (e22 = h23(e14)) & ~ (e22 = h23(e13)) & ~ (e22 = h23(e12)) & ~ (e22 = h23(e11)) & ~ (e22 = h23(e10))) | (~ (e21 = h23(e15)) & ~ (e21 = h23(e14)) & ~ (e21 = h23(e13)) & ~ (e21 = h23(e12)) & ~ (e21 = h23(e11)) & ~ (e21 = h23(e10))) | (~ (e20 = h23(e15)) & ~ (e20 = h23(e14)) & ~ (e20 = h23(e13)) & ~ (e20 = h23(e12)) & ~ (e20 = h23(e11)) & ~ (e20 = h23(e10))) | ~ (h23(op1(e15, e15)) = op2(h23(e15), h23(e15))) | ~ (h23(op1(e15, e14)) = op2(h23(e15), h23(e14))) | ~ (h23(op1(e15, e13)) = op2(h23(e15), h23(e13))) | ~ (h23(op1(e15, e12)) = op2(h23(e15), h23(e12))) | ~ (h23(op1(e15, e11)) = op2(h23(e15), h23(e11))) | ~ (h23(op1(e15, e10)) = op2(h23(e15), h23(e10))) | ~ (h23(op1(e14, e15)) = op2(h23(e14), h23(e15))) | ~ (h23(op1(e14, e14)) = op2(h23(e14), h23(e14))) | ~ (h23(op1(e14, e13)) = op2(h23(e14), h23(e13))) | ~ (h23(op1(e14, e12)) = op2(h23(e14), h23(e12))) | ~ (h23(op1(e14, e11)) = op2(h23(e14), h23(e11))) | ~ (h23(op1(e14, e10)) = op2(h23(e14), h23(e10))) | ~ (h23(op1(e13, e15)) = op2(h23(e13), h23(e15))) | ~ (h23(op1(e13, e14)) = op2(h23(e13), h23(e14))) | ~ (h23(op1(e13, e13)) = op2(h23(e13), h23(e13))) | ~ (h23(op1(e13, e12)) = op2(h23(e13), h23(e12))) | ~ (h23(op1(e13, e11)) = op2(h23(e13), h23(e11))) | ~ (h23(op1(e13, e10)) = op2(h23(e13), h23(e10))) | ~ (h23(op1(e12, e15)) = op2(h23(e12), h23(e15))) | ~ (h23(op1(e12, e14)) = op2(h23(e12), h23(e14))) | ~ (h23(op1(e12, e13)) = op2(h23(e12), h23(e13))) | ~ (h23(op1(e12, e12)) = op2(h23(e12), h23(e12))) | ~ (h23(op1(e12, e11)) = op2(h23(e12), h23(e11))) | ~ (h23(op1(e12, e10)) = op2(h23(e12), h23(e10))) | ~ (h23(op1(e11, e15)) = op2(h23(e11), h23(e15))) | ~ (h23(op1(e11, e14)) = op2(h23(e11), h23(e14))) | ~ (h23(op1(e11, e13)) = op2(h23(e11), h23(e13))) | ~ (h23(op1(e11, e12)) = op2(h23(e11), h23(e12))) | ~ (h23(op1(e11, e11)) = op2(h23(e11), h23(e11))) | ~ (h23(op1(e11, e10)) = op2(h23(e11), h23(e10))) | ~ (h23(op1(e10, e15)) = op2(h23(e10), h23(e15))) | ~ (h23(op1(e10, e14)) = op2(h23(e10), h23(e14))) | ~ (h23(op1(e10, e13)) = op2(h23(e10), h23(e13))) | ~ (h23(op1(e10, e12)) = op2(h23(e10), h23(e12))) | ~ (h23(op1(e10, e11)) = op2(h23(e10), h23(e11))) | ~ (h23(op1(e10, e10)) = op2(h23(e10), h23(e10)))) & ((~ (e25 = h22(e15)) & ~ (e25 = h22(e14)) & ~ (e25 = h22(e13)) & ~ (e25 = h22(e12)) & ~ (e25 = h22(e11)) & ~ (e25 = h22(e10))) | (~ (e24 = h22(e15)) & ~ (e24 = h22(e14)) & ~ (e24 = h22(e13)) & ~ (e24 = h22(e12)) & ~ (e24 = h22(e11)) & ~ (e24 = h22(e10))) | (~ (e23 = h22(e15)) & ~ (e23 = h22(e14)) & ~ (e23 = h22(e13)) & ~ (e23 = h22(e12)) & ~ (e23 = h22(e11)) & ~ (e23 = h22(e10))) | (~ (e22 = h22(e15)) & ~ (e22 = h22(e14)) & ~ (e22 = h22(e13)) & ~ (e22 = h22(e12)) & ~ (e22 = h22(e11)) & ~ (e22 = h22(e10))) | (~ (e21 = h22(e15)) & ~ (e21 = h22(e14)) & ~ (e21 = h22(e13)) & ~ (e21 = h22(e12)) & ~ (e21 = h22(e11)) & ~ (e21 = h22(e10))) | (~ (e20 = h22(e15)) & ~ (e20 = h22(e14)) & ~ (e20 = h22(e13)) & ~ (e20 = h22(e12)) & ~ (e20 = h22(e11)) & ~ (e20 = h22(e10))) | ~ (h22(op1(e15, e15)) = op2(h22(e15), h22(e15))) | ~ (h22(op1(e15, e14)) = op2(h22(e15), h22(e14))) | ~ (h22(op1(e15, e13)) = op2(h22(e15), h22(e13))) | ~ (h22(op1(e15, e12)) = op2(h22(e15), h22(e12))) | ~ (h22(op1(e15, e11)) = op2(h22(e15), h22(e11))) | ~ (h22(op1(e15, e10)) = op2(h22(e15), h22(e10))) | ~ (h22(op1(e14, e15)) = op2(h22(e14), h22(e15))) | ~ (h22(op1(e14, e14)) = op2(h22(e14), h22(e14))) | ~ (h22(op1(e14, e13)) = op2(h22(e14), h22(e13))) | ~ (h22(op1(e14, e12)) = op2(h22(e14), h22(e12))) | ~ (h22(op1(e14, e11)) = op2(h22(e14), h22(e11))) | ~ (h22(op1(e14, e10)) = op2(h22(e14), h22(e10))) | ~ (h22(op1(e13, e15)) = op2(h22(e13), h22(e15))) | ~ (h22(op1(e13, e14)) = op2(h22(e13), h22(e14))) | ~ (h22(op1(e13, e13)) = op2(h22(e13), h22(e13))) | ~ (h22(op1(e13, e12)) = op2(h22(e13), h22(e12))) | ~ (h22(op1(e13, e11)) = op2(h22(e13), h22(e11))) | ~ (h22(op1(e13, e10)) = op2(h22(e13), h22(e10))) | ~ (h22(op1(e12, e15)) = op2(h22(e12), h22(e15))) | ~ (h22(op1(e12, e14)) = op2(h22(e12), h22(e14))) | ~ (h22(op1(e12, e13)) = op2(h22(e12), h22(e13))) | ~ (h22(op1(e12, e12)) = op2(h22(e12), h22(e12))) | ~ (h22(op1(e12, e11)) = op2(h22(e12), h22(e11))) | ~ (h22(op1(e12, e10)) = op2(h22(e12), h22(e10))) | ~ (h22(op1(e11, e15)) = op2(h22(e11), h22(e15))) | ~ (h22(op1(e11, e14)) = op2(h22(e11), h22(e14))) | ~ (h22(op1(e11, e13)) = op2(h22(e11), h22(e13))) | ~ (h22(op1(e11, e12)) = op2(h22(e11), h22(e12))) | ~ (h22(op1(e11, e11)) = op2(h22(e11), h22(e11))) | ~ (h22(op1(e11, e10)) = op2(h22(e11), h22(e10))) | ~ (h22(op1(e10, e15)) = op2(h22(e10), h22(e15))) | ~ (h22(op1(e10, e14)) = op2(h22(e10), h22(e14))) | ~ (h22(op1(e10, e13)) = op2(h22(e10), h22(e13))) | ~ (h22(op1(e10, e12)) = op2(h22(e10), h22(e12))) | ~ (h22(op1(e10, e11)) = op2(h22(e10), h22(e11))) | ~ (h22(op1(e10, e10)) = op2(h22(e10), h22(e10)))) & ((~ (e25 = h21(e15)) & ~ (e25 = h21(e14)) & ~ (e25 = h21(e13)) & ~ (e25 = h21(e12)) & ~ (e25 = h21(e11)) & ~ (e25 = h21(e10))) | (~ (e24 = h21(e15)) & ~ (e24 = h21(e14)) & ~ (e24 = h21(e13)) & ~ (e24 = h21(e12)) & ~ (e24 = h21(e11)) & ~ (e24 = h21(e10))) | (~ (e23 = h21(e15)) & ~ (e23 = h21(e14)) & ~ (e23 = h21(e13)) & ~ (e23 = h21(e12)) & ~ (e23 = h21(e11)) & ~ (e23 = h21(e10))) | (~ (e22 = h21(e15)) & ~ (e22 = h21(e14)) & ~ (e22 = h21(e13)) & ~ (e22 = h21(e12)) & ~ (e22 = h21(e11)) & ~ (e22 = h21(e10))) | (~ (e21 = h21(e15)) & ~ (e21 = h21(e14)) & ~ (e21 = h21(e13)) & ~ (e21 = h21(e12)) & ~ (e21 = h21(e11)) & ~ (e21 = h21(e10))) | (~ (e20 = h21(e15)) & ~ (e20 = h21(e14)) & ~ (e20 = h21(e13)) & ~ (e20 = h21(e12)) & ~ (e20 = h21(e11)) & ~ (e20 = h21(e10))) | ~ (h21(op1(e15, e15)) = op2(h21(e15), h21(e15))) | ~ (h21(op1(e15, e14)) = op2(h21(e15), h21(e14))) | ~ (h21(op1(e15, e13)) = op2(h21(e15), h21(e13))) | ~ (h21(op1(e15, e12)) = op2(h21(e15), h21(e12))) | ~ (h21(op1(e15, e11)) = op2(h21(e15), h21(e11))) | ~ (h21(op1(e15, e10)) = op2(h21(e15), h21(e10))) | ~ (h21(op1(e14, e15)) = op2(h21(e14), h21(e15))) | ~ (h21(op1(e14, e14)) = op2(h21(e14), h21(e14))) | ~ (h21(op1(e14, e13)) = op2(h21(e14), h21(e13))) | ~ (h21(op1(e14, e12)) = op2(h21(e14), h21(e12))) | ~ (h21(op1(e14, e11)) = op2(h21(e14), h21(e11))) | ~ (h21(op1(e14, e10)) = op2(h21(e14), h21(e10))) | ~ (h21(op1(e13, e15)) = op2(h21(e13), h21(e15))) | ~ (h21(op1(e13, e14)) = op2(h21(e13), h21(e14))) | ~ (h21(op1(e13, e13)) = op2(h21(e13), h21(e13))) | ~ (h21(op1(e13, e12)) = op2(h21(e13), h21(e12))) | ~ (h21(op1(e13, e11)) = op2(h21(e13), h21(e11))) | ~ (h21(op1(e13, e10)) = op2(h21(e13), h21(e10))) | ~ (h21(op1(e12, e15)) = op2(h21(e12), h21(e15))) | ~ (h21(op1(e12, e14)) = op2(h21(e12), h21(e14))) | ~ (h21(op1(e12, e13)) = op2(h21(e12), h21(e13))) | ~ (h21(op1(e12, e12)) = op2(h21(e12), h21(e12))) | ~ (h21(op1(e12, e11)) = op2(h21(e12), h21(e11))) | ~ (h21(op1(e12, e10)) = op2(h21(e12), h21(e10))) | ~ (h21(op1(e11, e15)) = op2(h21(e11), h21(e15))) | ~ (h21(op1(e11, e14)) = op2(h21(e11), h21(e14))) | ~ (h21(op1(e11, e13)) = op2(h21(e11), h21(e13))) | ~ (h21(op1(e11, e12)) = op2(h21(e11), h21(e12))) | ~ (h21(op1(e11, e11)) = op2(h21(e11), h21(e11))) | ~ (h21(op1(e11, e10)) = op2(h21(e11), h21(e10))) | ~ (h21(op1(e10, e15)) = op2(h21(e10), h21(e15))) | ~ (h21(op1(e10, e14)) = op2(h21(e10), h21(e14))) | ~ (h21(op1(e10, e13)) = op2(h21(e10), h21(e13))) | ~ (h21(op1(e10, e12)) = op2(h21(e10), h21(e12))) | ~ (h21(op1(e10, e11)) = op2(h21(e10), h21(e11))) | ~ (h21(op1(e10, e10)) = op2(h21(e10), h21(e10)))) & ((~ (e25 = h20(e15)) & ~ (e25 = h20(e14)) & ~ (e25 = h20(e13)) & ~ (e25 = h20(e12)) & ~ (e25 = h20(e11)) & ~ (e25 = h20(e10))) | (~ (e24 = h20(e15)) & ~ (e24 = h20(e14)) & ~ (e24 = h20(e13)) & ~ (e24 = h20(e12)) & ~ (e24 = h20(e11)) & ~ (e24 = h20(e10))) | (~ (e23 = h20(e15)) & ~ (e23 = h20(e14)) & ~ (e23 = h20(e13)) & ~ (e23 = h20(e12)) & ~ (e23 = h20(e11)) & ~ (e23 = h20(e10))) | (~ (e22 = h20(e15)) & ~ (e22 = h20(e14)) & ~ (e22 = h20(e13)) & ~ (e22 = h20(e12)) & ~ (e22 = h20(e11)) & ~ (e22 = h20(e10))) | (~ (e21 = h20(e15)) & ~ (e21 = h20(e14)) & ~ (e21 = h20(e13)) & ~ (e21 = h20(e12)) & ~ (e21 = h20(e11)) & ~ (e21 = h20(e10))) | (~ (e20 = h20(e15)) & ~ (e20 = h20(e14)) & ~ (e20 = h20(e13)) & ~ (e20 = h20(e12)) & ~ (e20 = h20(e11)) & ~ (e20 = h20(e10))) | ~ (h20(op1(e15, e15)) = op2(h20(e15), h20(e15))) | ~ (h20(op1(e15, e14)) = op2(h20(e15), h20(e14))) | ~ (h20(op1(e15, e13)) = op2(h20(e15), h20(e13))) | ~ (h20(op1(e15, e12)) = op2(h20(e15), h20(e12))) | ~ (h20(op1(e15, e11)) = op2(h20(e15), h20(e11))) | ~ (h20(op1(e15, e10)) = op2(h20(e15), h20(e10))) | ~ (h20(op1(e14, e15)) = op2(h20(e14), h20(e15))) | ~ (h20(op1(e14, e14)) = op2(h20(e14), h20(e14))) | ~ (h20(op1(e14, e13)) = op2(h20(e14), h20(e13))) | ~ (h20(op1(e14, e12)) = op2(h20(e14), h20(e12))) | ~ (h20(op1(e14, e11)) = op2(h20(e14), h20(e11))) | ~ (h20(op1(e14, e10)) = op2(h20(e14), h20(e10))) | ~ (h20(op1(e13, e15)) = op2(h20(e13), h20(e15))) | ~ (h20(op1(e13, e14)) = op2(h20(e13), h20(e14))) | ~ (h20(op1(e13, e13)) = op2(h20(e13), h20(e13))) | ~ (h20(op1(e13, e12)) = op2(h20(e13), h20(e12))) | ~ (h20(op1(e13, e11)) = op2(h20(e13), h20(e11))) | ~ (h20(op1(e13, e10)) = op2(h20(e13), h20(e10))) | ~ (h20(op1(e12, e15)) = op2(h20(e12), h20(e15))) | ~ (h20(op1(e12, e14)) = op2(h20(e12), h20(e14))) | ~ (h20(op1(e12, e13)) = op2(h20(e12), h20(e13))) | ~ (h20(op1(e12, e12)) = op2(h20(e12), h20(e12))) | ~ (h20(op1(e12, e11)) = op2(h20(e12), h20(e11))) | ~ (h20(op1(e12, e10)) = op2(h20(e12), h20(e10))) | ~ (h20(op1(e11, e15)) = op2(h20(e11), h20(e15))) | ~ (h20(op1(e11, e14)) = op2(h20(e11), h20(e14))) | ~ (h20(op1(e11, e13)) = op2(h20(e11), h20(e13))) | ~ (h20(op1(e11, e12)) = op2(h20(e11), h20(e12))) | ~ (h20(op1(e11, e11)) = op2(h20(e11), h20(e11))) | ~ (h20(op1(e11, e10)) = op2(h20(e11), h20(e10))) | ~ (h20(op1(e10, e15)) = op2(h20(e10), h20(e15))) | ~ (h20(op1(e10, e14)) = op2(h20(e10), h20(e14))) | ~ (h20(op1(e10, e13)) = op2(h20(e10), h20(e13))) | ~ (h20(op1(e10, e12)) = op2(h20(e10), h20(e12))) | ~ (h20(op1(e10, e11)) = op2(h20(e10), h20(e11))) | ~ (h20(op1(e10, e10)) = op2(h20(e10), h20(e10)))) & ((~ (e25 = h19(e15)) & ~ (e25 = h19(e14)) & ~ (e25 = h19(e13)) & ~ (e25 = h19(e12)) & ~ (e25 = h19(e11)) & ~ (e25 = h19(e10))) | (~ (e24 = h19(e15)) & ~ (e24 = h19(e14)) & ~ (e24 = h19(e13)) & ~ (e24 = h19(e12)) & ~ (e24 = h19(e11)) & ~ (e24 = h19(e10))) | (~ (e23 = h19(e15)) & ~ (e23 = h19(e14)) & ~ (e23 = h19(e13)) & ~ (e23 = h19(e12)) & ~ (e23 = h19(e11)) & ~ (e23 = h19(e10))) | (~ (e22 = h19(e15)) & ~ (e22 = h19(e14)) & ~ (e22 = h19(e13)) & ~ (e22 = h19(e12)) & ~ (e22 = h19(e11)) & ~ (e22 = h19(e10))) | (~ (e21 = h19(e15)) & ~ (e21 = h19(e14)) & ~ (e21 = h19(e13)) & ~ (e21 = h19(e12)) & ~ (e21 = h19(e11)) & ~ (e21 = h19(e10))) | (~ (e20 = h19(e15)) & ~ (e20 = h19(e14)) & ~ (e20 = h19(e13)) & ~ (e20 = h19(e12)) & ~ (e20 = h19(e11)) & ~ (e20 = h19(e10))) | ~ (h19(op1(e15, e15)) = op2(h19(e15), h19(e15))) | ~ (h19(op1(e15, e14)) = op2(h19(e15), h19(e14))) | ~ (h19(op1(e15, e13)) = op2(h19(e15), h19(e13))) | ~ (h19(op1(e15, e12)) = op2(h19(e15), h19(e12))) | ~ (h19(op1(e15, e11)) = op2(h19(e15), h19(e11))) | ~ (h19(op1(e15, e10)) = op2(h19(e15), h19(e10))) | ~ (h19(op1(e14, e15)) = op2(h19(e14), h19(e15))) | ~ (h19(op1(e14, e14)) = op2(h19(e14), h19(e14))) | ~ (h19(op1(e14, e13)) = op2(h19(e14), h19(e13))) | ~ (h19(op1(e14, e12)) = op2(h19(e14), h19(e12))) | ~ (h19(op1(e14, e11)) = op2(h19(e14), h19(e11))) | ~ (h19(op1(e14, e10)) = op2(h19(e14), h19(e10))) | ~ (h19(op1(e13, e15)) = op2(h19(e13), h19(e15))) | ~ (h19(op1(e13, e14)) = op2(h19(e13), h19(e14))) | ~ (h19(op1(e13, e13)) = op2(h19(e13), h19(e13))) | ~ (h19(op1(e13, e12)) = op2(h19(e13), h19(e12))) | ~ (h19(op1(e13, e11)) = op2(h19(e13), h19(e11))) | ~ (h19(op1(e13, e10)) = op2(h19(e13), h19(e10))) | ~ (h19(op1(e12, e15)) = op2(h19(e12), h19(e15))) | ~ (h19(op1(e12, e14)) = op2(h19(e12), h19(e14))) | ~ (h19(op1(e12, e13)) = op2(h19(e12), h19(e13))) | ~ (h19(op1(e12, e12)) = op2(h19(e12), h19(e12))) | ~ (h19(op1(e12, e11)) = op2(h19(e12), h19(e11))) | ~ (h19(op1(e12, e10)) = op2(h19(e12), h19(e10))) | ~ (h19(op1(e11, e15)) = op2(h19(e11), h19(e15))) | ~ (h19(op1(e11, e14)) = op2(h19(e11), h19(e14))) | ~ (h19(op1(e11, e13)) = op2(h19(e11), h19(e13))) | ~ (h19(op1(e11, e12)) = op2(h19(e11), h19(e12))) | ~ (h19(op1(e11, e11)) = op2(h19(e11), h19(e11))) | ~ (h19(op1(e11, e10)) = op2(h19(e11), h19(e10))) | ~ (h19(op1(e10, e15)) = op2(h19(e10), h19(e15))) | ~ (h19(op1(e10, e14)) = op2(h19(e10), h19(e14))) | ~ (h19(op1(e10, e13)) = op2(h19(e10), h19(e13))) | ~ (h19(op1(e10, e12)) = op2(h19(e10), h19(e12))) | ~ (h19(op1(e10, e11)) = op2(h19(e10), h19(e11))) | ~ (h19(op1(e10, e10)) = op2(h19(e10), h19(e10)))) & ((~ (e25 = h18(e15)) & ~ (e25 = h18(e14)) & ~ (e25 = h18(e13)) & ~ (e25 = h18(e12)) & ~ (e25 = h18(e11)) & ~ (e25 = h18(e10))) | (~ (e24 = h18(e15)) & ~ (e24 = h18(e14)) & ~ (e24 = h18(e13)) & ~ (e24 = h18(e12)) & ~ (e24 = h18(e11)) & ~ (e24 = h18(e10))) | (~ (e23 = h18(e15)) & ~ (e23 = h18(e14)) & ~ (e23 = h18(e13)) & ~ (e23 = h18(e12)) & ~ (e23 = h18(e11)) & ~ (e23 = h18(e10))) | (~ (e22 = h18(e15)) & ~ (e22 = h18(e14)) & ~ (e22 = h18(e13)) & ~ (e22 = h18(e12)) & ~ (e22 = h18(e11)) & ~ (e22 = h18(e10))) | (~ (e21 = h18(e15)) & ~ (e21 = h18(e14)) & ~ (e21 = h18(e13)) & ~ (e21 = h18(e12)) & ~ (e21 = h18(e11)) & ~ (e21 = h18(e10))) | (~ (e20 = h18(e15)) & ~ (e20 = h18(e14)) & ~ (e20 = h18(e13)) & ~ (e20 = h18(e12)) & ~ (e20 = h18(e11)) & ~ (e20 = h18(e10))) | ~ (h18(op1(e15, e15)) = op2(h18(e15), h18(e15))) | ~ (h18(op1(e15, e14)) = op2(h18(e15), h18(e14))) | ~ (h18(op1(e15, e13)) = op2(h18(e15), h18(e13))) | ~ (h18(op1(e15, e12)) = op2(h18(e15), h18(e12))) | ~ (h18(op1(e15, e11)) = op2(h18(e15), h18(e11))) | ~ (h18(op1(e15, e10)) = op2(h18(e15), h18(e10))) | ~ (h18(op1(e14, e15)) = op2(h18(e14), h18(e15))) | ~ (h18(op1(e14, e14)) = op2(h18(e14), h18(e14))) | ~ (h18(op1(e14, e13)) = op2(h18(e14), h18(e13))) | ~ (h18(op1(e14, e12)) = op2(h18(e14), h18(e12))) | ~ (h18(op1(e14, e11)) = op2(h18(e14), h18(e11))) | ~ (h18(op1(e14, e10)) = op2(h18(e14), h18(e10))) | ~ (h18(op1(e13, e15)) = op2(h18(e13), h18(e15))) | ~ (h18(op1(e13, e14)) = op2(h18(e13), h18(e14))) | ~ (h18(op1(e13, e13)) = op2(h18(e13), h18(e13))) | ~ (h18(op1(e13, e12)) = op2(h18(e13), h18(e12))) | ~ (h18(op1(e13, e11)) = op2(h18(e13), h18(e11))) | ~ (h18(op1(e13, e10)) = op2(h18(e13), h18(e10))) | ~ (h18(op1(e12, e15)) = op2(h18(e12), h18(e15))) | ~ (h18(op1(e12, e14)) = op2(h18(e12), h18(e14))) | ~ (h18(op1(e12, e13)) = op2(h18(e12), h18(e13))) | ~ (h18(op1(e12, e12)) = op2(h18(e12), h18(e12))) | ~ (h18(op1(e12, e11)) = op2(h18(e12), h18(e11))) | ~ (h18(op1(e12, e10)) = op2(h18(e12), h18(e10))) | ~ (h18(op1(e11, e15)) = op2(h18(e11), h18(e15))) | ~ (h18(op1(e11, e14)) = op2(h18(e11), h18(e14))) | ~ (h18(op1(e11, e13)) = op2(h18(e11), h18(e13))) | ~ (h18(op1(e11, e12)) = op2(h18(e11), h18(e12))) | ~ (h18(op1(e11, e11)) = op2(h18(e11), h18(e11))) | ~ (h18(op1(e11, e10)) = op2(h18(e11), h18(e10))) | ~ (h18(op1(e10, e15)) = op2(h18(e10), h18(e15))) | ~ (h18(op1(e10, e14)) = op2(h18(e10), h18(e14))) | ~ (h18(op1(e10, e13)) = op2(h18(e10), h18(e13))) | ~ (h18(op1(e10, e12)) = op2(h18(e10), h18(e12))) | ~ (h18(op1(e10, e11)) = op2(h18(e10), h18(e11))) | ~ (h18(op1(e10, e10)) = op2(h18(e10), h18(e10)))) & ((~ (e25 = h17(e15)) & ~ (e25 = h17(e14)) & ~ (e25 = h17(e13)) & ~ (e25 = h17(e12)) & ~ (e25 = h17(e11)) & ~ (e25 = h17(e10))) | (~ (e24 = h17(e15)) & ~ (e24 = h17(e14)) & ~ (e24 = h17(e13)) & ~ (e24 = h17(e12)) & ~ (e24 = h17(e11)) & ~ (e24 = h17(e10))) | (~ (e23 = h17(e15)) & ~ (e23 = h17(e14)) & ~ (e23 = h17(e13)) & ~ (e23 = h17(e12)) & ~ (e23 = h17(e11)) & ~ (e23 = h17(e10))) | (~ (e22 = h17(e15)) & ~ (e22 = h17(e14)) & ~ (e22 = h17(e13)) & ~ (e22 = h17(e12)) & ~ (e22 = h17(e11)) & ~ (e22 = h17(e10))) | (~ (e21 = h17(e15)) & ~ (e21 = h17(e14)) & ~ (e21 = h17(e13)) & ~ (e21 = h17(e12)) & ~ (e21 = h17(e11)) & ~ (e21 = h17(e10))) | (~ (e20 = h17(e15)) & ~ (e20 = h17(e14)) & ~ (e20 = h17(e13)) & ~ (e20 = h17(e12)) & ~ (e20 = h17(e11)) & ~ (e20 = h17(e10))) | ~ (h17(op1(e15, e15)) = op2(h17(e15), h17(e15))) | ~ (h17(op1(e15, e14)) = op2(h17(e15), h17(e14))) | ~ (h17(op1(e15, e13)) = op2(h17(e15), h17(e13))) | ~ (h17(op1(e15, e12)) = op2(h17(e15), h17(e12))) | ~ (h17(op1(e15, e11)) = op2(h17(e15), h17(e11))) | ~ (h17(op1(e15, e10)) = op2(h17(e15), h17(e10))) | ~ (h17(op1(e14, e15)) = op2(h17(e14), h17(e15))) | ~ (h17(op1(e14, e14)) = op2(h17(e14), h17(e14))) | ~ (h17(op1(e14, e13)) = op2(h17(e14), h17(e13))) | ~ (h17(op1(e14, e12)) = op2(h17(e14), h17(e12))) | ~ (h17(op1(e14, e11)) = op2(h17(e14), h17(e11))) | ~ (h17(op1(e14, e10)) = op2(h17(e14), h17(e10))) | ~ (h17(op1(e13, e15)) = op2(h17(e13), h17(e15))) | ~ (h17(op1(e13, e14)) = op2(h17(e13), h17(e14))) | ~ (h17(op1(e13, e13)) = op2(h17(e13), h17(e13))) | ~ (h17(op1(e13, e12)) = op2(h17(e13), h17(e12))) | ~ (h17(op1(e13, e11)) = op2(h17(e13), h17(e11))) | ~ (h17(op1(e13, e10)) = op2(h17(e13), h17(e10))) | ~ (h17(op1(e12, e15)) = op2(h17(e12), h17(e15))) | ~ (h17(op1(e12, e14)) = op2(h17(e12), h17(e14))) | ~ (h17(op1(e12, e13)) = op2(h17(e12), h17(e13))) | ~ (h17(op1(e12, e12)) = op2(h17(e12), h17(e12))) | ~ (h17(op1(e12, e11)) = op2(h17(e12), h17(e11))) | ~ (h17(op1(e12, e10)) = op2(h17(e12), h17(e10))) | ~ (h17(op1(e11, e15)) = op2(h17(e11), h17(e15))) | ~ (h17(op1(e11, e14)) = op2(h17(e11), h17(e14))) | ~ (h17(op1(e11, e13)) = op2(h17(e11), h17(e13))) | ~ (h17(op1(e11, e12)) = op2(h17(e11), h17(e12))) | ~ (h17(op1(e11, e11)) = op2(h17(e11), h17(e11))) | ~ (h17(op1(e11, e10)) = op2(h17(e11), h17(e10))) | ~ (h17(op1(e10, e15)) = op2(h17(e10), h17(e15))) | ~ (h17(op1(e10, e14)) = op2(h17(e10), h17(e14))) | ~ (h17(op1(e10, e13)) = op2(h17(e10), h17(e13))) | ~ (h17(op1(e10, e12)) = op2(h17(e10), h17(e12))) | ~ (h17(op1(e10, e11)) = op2(h17(e10), h17(e11))) | ~ (h17(op1(e10, e10)) = op2(h17(e10), h17(e10)))) & ((~ (e25 = h16(e15)) & ~ (e25 = h16(e14)) & ~ (e25 = h16(e13)) & ~ (e25 = h16(e12)) & ~ (e25 = h16(e11)) & ~ (e25 = h16(e10))) | (~ (e24 = h16(e15)) & ~ (e24 = h16(e14)) & ~ (e24 = h16(e13)) & ~ (e24 = h16(e12)) & ~ (e24 = h16(e11)) & ~ (e24 = h16(e10))) | (~ (e23 = h16(e15)) & ~ (e23 = h16(e14)) & ~ (e23 = h16(e13)) & ~ (e23 = h16(e12)) & ~ (e23 = h16(e11)) & ~ (e23 = h16(e10))) | (~ (e22 = h16(e15)) & ~ (e22 = h16(e14)) & ~ (e22 = h16(e13)) & ~ (e22 = h16(e12)) & ~ (e22 = h16(e11)) & ~ (e22 = h16(e10))) | (~ (e21 = h16(e15)) & ~ (e21 = h16(e14)) & ~ (e21 = h16(e13)) & ~ (e21 = h16(e12)) & ~ (e21 = h16(e11)) & ~ (e21 = h16(e10))) | (~ (e20 = h16(e15)) & ~ (e20 = h16(e14)) & ~ (e20 = h16(e13)) & ~ (e20 = h16(e12)) & ~ (e20 = h16(e11)) & ~ (e20 = h16(e10))) | ~ (h16(op1(e15, e15)) = op2(h16(e15), h16(e15))) | ~ (h16(op1(e15, e14)) = op2(h16(e15), h16(e14))) | ~ (h16(op1(e15, e13)) = op2(h16(e15), h16(e13))) | ~ (h16(op1(e15, e12)) = op2(h16(e15), h16(e12))) | ~ (h16(op1(e15, e11)) = op2(h16(e15), h16(e11))) | ~ (h16(op1(e15, e10)) = op2(h16(e15), h16(e10))) | ~ (h16(op1(e14, e15)) = op2(h16(e14), h16(e15))) | ~ (h16(op1(e14, e14)) = op2(h16(e14), h16(e14))) | ~ (h16(op1(e14, e13)) = op2(h16(e14), h16(e13))) | ~ (h16(op1(e14, e12)) = op2(h16(e14), h16(e12))) | ~ (h16(op1(e14, e11)) = op2(h16(e14), h16(e11))) | ~ (h16(op1(e14, e10)) = op2(h16(e14), h16(e10))) | ~ (h16(op1(e13, e15)) = op2(h16(e13), h16(e15))) | ~ (h16(op1(e13, e14)) = op2(h16(e13), h16(e14))) | ~ (h16(op1(e13, e13)) = op2(h16(e13), h16(e13))) | ~ (h16(op1(e13, e12)) = op2(h16(e13), h16(e12))) | ~ (h16(op1(e13, e11)) = op2(h16(e13), h16(e11))) | ~ (h16(op1(e13, e10)) = op2(h16(e13), h16(e10))) | ~ (h16(op1(e12, e15)) = op2(h16(e12), h16(e15))) | ~ (h16(op1(e12, e14)) = op2(h16(e12), h16(e14))) | ~ (h16(op1(e12, e13)) = op2(h16(e12), h16(e13))) | ~ (h16(op1(e12, e12)) = op2(h16(e12), h16(e12))) | ~ (h16(op1(e12, e11)) = op2(h16(e12), h16(e11))) | ~ (h16(op1(e12, e10)) = op2(h16(e12), h16(e10))) | ~ (h16(op1(e11, e15)) = op2(h16(e11), h16(e15))) | ~ (h16(op1(e11, e14)) = op2(h16(e11), h16(e14))) | ~ (h16(op1(e11, e13)) = op2(h16(e11), h16(e13))) | ~ (h16(op1(e11, e12)) = op2(h16(e11), h16(e12))) | ~ (h16(op1(e11, e11)) = op2(h16(e11), h16(e11))) | ~ (h16(op1(e11, e10)) = op2(h16(e11), h16(e10))) | ~ (h16(op1(e10, e15)) = op2(h16(e10), h16(e15))) | ~ (h16(op1(e10, e14)) = op2(h16(e10), h16(e14))) | ~ (h16(op1(e10, e13)) = op2(h16(e10), h16(e13))) | ~ (h16(op1(e10, e12)) = op2(h16(e10), h16(e12))) | ~ (h16(op1(e10, e11)) = op2(h16(e10), h16(e11))) | ~ (h16(op1(e10, e10)) = op2(h16(e10), h16(e10)))) & ((~ (e25 = h15(e15)) & ~ (e25 = h15(e14)) & ~ (e25 = h15(e13)) & ~ (e25 = h15(e12)) & ~ (e25 = h15(e11)) & ~ (e25 = h15(e10))) | (~ (e24 = h15(e15)) & ~ (e24 = h15(e14)) & ~ (e24 = h15(e13)) & ~ (e24 = h15(e12)) & ~ (e24 = h15(e11)) & ~ (e24 = h15(e10))) | (~ (e23 = h15(e15)) & ~ (e23 = h15(e14)) & ~ (e23 = h15(e13)) & ~ (e23 = h15(e12)) & ~ (e23 = h15(e11)) & ~ (e23 = h15(e10))) | (~ (e22 = h15(e15)) & ~ (e22 = h15(e14)) & ~ (e22 = h15(e13)) & ~ (e22 = h15(e12)) & ~ (e22 = h15(e11)) & ~ (e22 = h15(e10))) | (~ (e21 = h15(e15)) & ~ (e21 = h15(e14)) & ~ (e21 = h15(e13)) & ~ (e21 = h15(e12)) & ~ (e21 = h15(e11)) & ~ (e21 = h15(e10))) | (~ (e20 = h15(e15)) & ~ (e20 = h15(e14)) & ~ (e20 = h15(e13)) & ~ (e20 = h15(e12)) & ~ (e20 = h15(e11)) & ~ (e20 = h15(e10))) | ~ (h15(op1(e15, e15)) = op2(h15(e15), h15(e15))) | ~ (h15(op1(e15, e14)) = op2(h15(e15), h15(e14))) | ~ (h15(op1(e15, e13)) = op2(h15(e15), h15(e13))) | ~ (h15(op1(e15, e12)) = op2(h15(e15), h15(e12))) | ~ (h15(op1(e15, e11)) = op2(h15(e15), h15(e11))) | ~ (h15(op1(e15, e10)) = op2(h15(e15), h15(e10))) | ~ (h15(op1(e14, e15)) = op2(h15(e14), h15(e15))) | ~ (h15(op1(e14, e14)) = op2(h15(e14), h15(e14))) | ~ (h15(op1(e14, e13)) = op2(h15(e14), h15(e13))) | ~ (h15(op1(e14, e12)) = op2(h15(e14), h15(e12))) | ~ (h15(op1(e14, e11)) = op2(h15(e14), h15(e11))) | ~ (h15(op1(e14, e10)) = op2(h15(e14), h15(e10))) | ~ (h15(op1(e13, e15)) = op2(h15(e13), h15(e15))) | ~ (h15(op1(e13, e14)) = op2(h15(e13), h15(e14))) | ~ (h15(op1(e13, e13)) = op2(h15(e13), h15(e13))) | ~ (h15(op1(e13, e12)) = op2(h15(e13), h15(e12))) | ~ (h15(op1(e13, e11)) = op2(h15(e13), h15(e11))) | ~ (h15(op1(e13, e10)) = op2(h15(e13), h15(e10))) | ~ (h15(op1(e12, e15)) = op2(h15(e12), h15(e15))) | ~ (h15(op1(e12, e14)) = op2(h15(e12), h15(e14))) | ~ (h15(op1(e12, e13)) = op2(h15(e12), h15(e13))) | ~ (h15(op1(e12, e12)) = op2(h15(e12), h15(e12))) | ~ (h15(op1(e12, e11)) = op2(h15(e12), h15(e11))) | ~ (h15(op1(e12, e10)) = op2(h15(e12), h15(e10))) | ~ (h15(op1(e11, e15)) = op2(h15(e11), h15(e15))) | ~ (h15(op1(e11, e14)) = op2(h15(e11), h15(e14))) | ~ (h15(op1(e11, e13)) = op2(h15(e11), h15(e13))) | ~ (h15(op1(e11, e12)) = op2(h15(e11), h15(e12))) | ~ (h15(op1(e11, e11)) = op2(h15(e11), h15(e11))) | ~ (h15(op1(e11, e10)) = op2(h15(e11), h15(e10))) | ~ (h15(op1(e10, e15)) = op2(h15(e10), h15(e15))) | ~ (h15(op1(e10, e14)) = op2(h15(e10), h15(e14))) | ~ (h15(op1(e10, e13)) = op2(h15(e10), h15(e13))) | ~ (h15(op1(e10, e12)) = op2(h15(e10), h15(e12))) | ~ (h15(op1(e10, e11)) = op2(h15(e10), h15(e11))) | ~ (h15(op1(e10, e10)) = op2(h15(e10), h15(e10)))) & ((~ (e25 = h14(e15)) & ~ (e25 = h14(e14)) & ~ (e25 = h14(e13)) & ~ (e25 = h14(e12)) & ~ (e25 = h14(e11)) & ~ (e25 = h14(e10))) | (~ (e24 = h14(e15)) & ~ (e24 = h14(e14)) & ~ (e24 = h14(e13)) & ~ (e24 = h14(e12)) & ~ (e24 = h14(e11)) & ~ (e24 = h14(e10))) | (~ (e23 = h14(e15)) & ~ (e23 = h14(e14)) & ~ (e23 = h14(e13)) & ~ (e23 = h14(e12)) & ~ (e23 = h14(e11)) & ~ (e23 = h14(e10))) | (~ (e22 = h14(e15)) & ~ (e22 = h14(e14)) & ~ (e22 = h14(e13)) & ~ (e22 = h14(e12)) & ~ (e22 = h14(e11)) & ~ (e22 = h14(e10))) | (~ (e21 = h14(e15)) & ~ (e21 = h14(e14)) & ~ (e21 = h14(e13)) & ~ (e21 = h14(e12)) & ~ (e21 = h14(e11)) & ~ (e21 = h14(e10))) | (~ (e20 = h14(e15)) & ~ (e20 = h14(e14)) & ~ (e20 = h14(e13)) & ~ (e20 = h14(e12)) & ~ (e20 = h14(e11)) & ~ (e20 = h14(e10))) | ~ (h14(op1(e15, e15)) = op2(h14(e15), h14(e15))) | ~ (h14(op1(e15, e14)) = op2(h14(e15), h14(e14))) | ~ (h14(op1(e15, e13)) = op2(h14(e15), h14(e13))) | ~ (h14(op1(e15, e12)) = op2(h14(e15), h14(e12))) | ~ (h14(op1(e15, e11)) = op2(h14(e15), h14(e11))) | ~ (h14(op1(e15, e10)) = op2(h14(e15), h14(e10))) | ~ (h14(op1(e14, e15)) = op2(h14(e14), h14(e15))) | ~ (h14(op1(e14, e14)) = op2(h14(e14), h14(e14))) | ~ (h14(op1(e14, e13)) = op2(h14(e14), h14(e13))) | ~ (h14(op1(e14, e12)) = op2(h14(e14), h14(e12))) | ~ (h14(op1(e14, e11)) = op2(h14(e14), h14(e11))) | ~ (h14(op1(e14, e10)) = op2(h14(e14), h14(e10))) | ~ (h14(op1(e13, e15)) = op2(h14(e13), h14(e15))) | ~ (h14(op1(e13, e14)) = op2(h14(e13), h14(e14))) | ~ (h14(op1(e13, e13)) = op2(h14(e13), h14(e13))) | ~ (h14(op1(e13, e12)) = op2(h14(e13), h14(e12))) | ~ (h14(op1(e13, e11)) = op2(h14(e13), h14(e11))) | ~ (h14(op1(e13, e10)) = op2(h14(e13), h14(e10))) | ~ (h14(op1(e12, e15)) = op2(h14(e12), h14(e15))) | ~ (h14(op1(e12, e14)) = op2(h14(e12), h14(e14))) | ~ (h14(op1(e12, e13)) = op2(h14(e12), h14(e13))) | ~ (h14(op1(e12, e12)) = op2(h14(e12), h14(e12))) | ~ (h14(op1(e12, e11)) = op2(h14(e12), h14(e11))) | ~ (h14(op1(e12, e10)) = op2(h14(e12), h14(e10))) | ~ (h14(op1(e11, e15)) = op2(h14(e11), h14(e15))) | ~ (h14(op1(e11, e14)) = op2(h14(e11), h14(e14))) | ~ (h14(op1(e11, e13)) = op2(h14(e11), h14(e13))) | ~ (h14(op1(e11, e12)) = op2(h14(e11), h14(e12))) | ~ (h14(op1(e11, e11)) = op2(h14(e11), h14(e11))) | ~ (h14(op1(e11, e10)) = op2(h14(e11), h14(e10))) | ~ (h14(op1(e10, e15)) = op2(h14(e10), h14(e15))) | ~ (h14(op1(e10, e14)) = op2(h14(e10), h14(e14))) | ~ (h14(op1(e10, e13)) = op2(h14(e10), h14(e13))) | ~ (h14(op1(e10, e12)) = op2(h14(e10), h14(e12))) | ~ (h14(op1(e10, e11)) = op2(h14(e10), h14(e11))) | ~ (h14(op1(e10, e10)) = op2(h14(e10), h14(e10)))) & ((~ (e25 = h13(e15)) & ~ (e25 = h13(e14)) & ~ (e25 = h13(e13)) & ~ (e25 = h13(e12)) & ~ (e25 = h13(e11)) & ~ (e25 = h13(e10))) | (~ (e24 = h13(e15)) & ~ (e24 = h13(e14)) & ~ (e24 = h13(e13)) & ~ (e24 = h13(e12)) & ~ (e24 = h13(e11)) & ~ (e24 = h13(e10))) | (~ (e23 = h13(e15)) & ~ (e23 = h13(e14)) & ~ (e23 = h13(e13)) & ~ (e23 = h13(e12)) & ~ (e23 = h13(e11)) & ~ (e23 = h13(e10))) | (~ (e22 = h13(e15)) & ~ (e22 = h13(e14)) & ~ (e22 = h13(e13)) & ~ (e22 = h13(e12)) & ~ (e22 = h13(e11)) & ~ (e22 = h13(e10))) | (~ (e21 = h13(e15)) & ~ (e21 = h13(e14)) & ~ (e21 = h13(e13)) & ~ (e21 = h13(e12)) & ~ (e21 = h13(e11)) & ~ (e21 = h13(e10))) | (~ (e20 = h13(e15)) & ~ (e20 = h13(e14)) & ~ (e20 = h13(e13)) & ~ (e20 = h13(e12)) & ~ (e20 = h13(e11)) & ~ (e20 = h13(e10))) | ~ (h13(op1(e15, e15)) = op2(h13(e15), h13(e15))) | ~ (h13(op1(e15, e14)) = op2(h13(e15), h13(e14))) | ~ (h13(op1(e15, e13)) = op2(h13(e15), h13(e13))) | ~ (h13(op1(e15, e12)) = op2(h13(e15), h13(e12))) | ~ (h13(op1(e15, e11)) = op2(h13(e15), h13(e11))) | ~ (h13(op1(e15, e10)) = op2(h13(e15), h13(e10))) | ~ (h13(op1(e14, e15)) = op2(h13(e14), h13(e15))) | ~ (h13(op1(e14, e14)) = op2(h13(e14), h13(e14))) | ~ (h13(op1(e14, e13)) = op2(h13(e14), h13(e13))) | ~ (h13(op1(e14, e12)) = op2(h13(e14), h13(e12))) | ~ (h13(op1(e14, e11)) = op2(h13(e14), h13(e11))) | ~ (h13(op1(e14, e10)) = op2(h13(e14), h13(e10))) | ~ (h13(op1(e13, e15)) = op2(h13(e13), h13(e15))) | ~ (h13(op1(e13, e14)) = op2(h13(e13), h13(e14))) | ~ (h13(op1(e13, e13)) = op2(h13(e13), h13(e13))) | ~ (h13(op1(e13, e12)) = op2(h13(e13), h13(e12))) | ~ (h13(op1(e13, e11)) = op2(h13(e13), h13(e11))) | ~ (h13(op1(e13, e10)) = op2(h13(e13), h13(e10))) | ~ (h13(op1(e12, e15)) = op2(h13(e12), h13(e15))) | ~ (h13(op1(e12, e14)) = op2(h13(e12), h13(e14))) | ~ (h13(op1(e12, e13)) = op2(h13(e12), h13(e13))) | ~ (h13(op1(e12, e12)) = op2(h13(e12), h13(e12))) | ~ (h13(op1(e12, e11)) = op2(h13(e12), h13(e11))) | ~ (h13(op1(e12, e10)) = op2(h13(e12), h13(e10))) | ~ (h13(op1(e11, e15)) = op2(h13(e11), h13(e15))) | ~ (h13(op1(e11, e14)) = op2(h13(e11), h13(e14))) | ~ (h13(op1(e11, e13)) = op2(h13(e11), h13(e13))) | ~ (h13(op1(e11, e12)) = op2(h13(e11), h13(e12))) | ~ (h13(op1(e11, e11)) = op2(h13(e11), h13(e11))) | ~ (h13(op1(e11, e10)) = op2(h13(e11), h13(e10))) | ~ (h13(op1(e10, e15)) = op2(h13(e10), h13(e15))) | ~ (h13(op1(e10, e14)) = op2(h13(e10), h13(e14))) | ~ (h13(op1(e10, e13)) = op2(h13(e10), h13(e13))) | ~ (h13(op1(e10, e12)) = op2(h13(e10), h13(e12))) | ~ (h13(op1(e10, e11)) = op2(h13(e10), h13(e11))) | ~ (h13(op1(e10, e10)) = op2(h13(e10), h13(e10)))) & ((~ (e25 = h12(e15)) & ~ (e25 = h12(e14)) & ~ (e25 = h12(e13)) & ~ (e25 = h12(e12)) & ~ (e25 = h12(e11)) & ~ (e25 = h12(e10))) | (~ (e24 = h12(e15)) & ~ (e24 = h12(e14)) & ~ (e24 = h12(e13)) & ~ (e24 = h12(e12)) & ~ (e24 = h12(e11)) & ~ (e24 = h12(e10))) | (~ (e23 = h12(e15)) & ~ (e23 = h12(e14)) & ~ (e23 = h12(e13)) & ~ (e23 = h12(e12)) & ~ (e23 = h12(e11)) & ~ (e23 = h12(e10))) | (~ (e22 = h12(e15)) & ~ (e22 = h12(e14)) & ~ (e22 = h12(e13)) & ~ (e22 = h12(e12)) & ~ (e22 = h12(e11)) & ~ (e22 = h12(e10))) | (~ (e21 = h12(e15)) & ~ (e21 = h12(e14)) & ~ (e21 = h12(e13)) & ~ (e21 = h12(e12)) & ~ (e21 = h12(e11)) & ~ (e21 = h12(e10))) | (~ (e20 = h12(e15)) & ~ (e20 = h12(e14)) & ~ (e20 = h12(e13)) & ~ (e20 = h12(e12)) & ~ (e20 = h12(e11)) & ~ (e20 = h12(e10))) | ~ (h12(op1(e15, e15)) = op2(h12(e15), h12(e15))) | ~ (h12(op1(e15, e14)) = op2(h12(e15), h12(e14))) | ~ (h12(op1(e15, e13)) = op2(h12(e15), h12(e13))) | ~ (h12(op1(e15, e12)) = op2(h12(e15), h12(e12))) | ~ (h12(op1(e15, e11)) = op2(h12(e15), h12(e11))) | ~ (h12(op1(e15, e10)) = op2(h12(e15), h12(e10))) | ~ (h12(op1(e14, e15)) = op2(h12(e14), h12(e15))) | ~ (h12(op1(e14, e14)) = op2(h12(e14), h12(e14))) | ~ (h12(op1(e14, e13)) = op2(h12(e14), h12(e13))) | ~ (h12(op1(e14, e12)) = op2(h12(e14), h12(e12))) | ~ (h12(op1(e14, e11)) = op2(h12(e14), h12(e11))) | ~ (h12(op1(e14, e10)) = op2(h12(e14), h12(e10))) | ~ (h12(op1(e13, e15)) = op2(h12(e13), h12(e15))) | ~ (h12(op1(e13, e14)) = op2(h12(e13), h12(e14))) | ~ (h12(op1(e13, e13)) = op2(h12(e13), h12(e13))) | ~ (h12(op1(e13, e12)) = op2(h12(e13), h12(e12))) | ~ (h12(op1(e13, e11)) = op2(h12(e13), h12(e11))) | ~ (h12(op1(e13, e10)) = op2(h12(e13), h12(e10))) | ~ (h12(op1(e12, e15)) = op2(h12(e12), h12(e15))) | ~ (h12(op1(e12, e14)) = op2(h12(e12), h12(e14))) | ~ (h12(op1(e12, e13)) = op2(h12(e12), h12(e13))) | ~ (h12(op1(e12, e12)) = op2(h12(e12), h12(e12))) | ~ (h12(op1(e12, e11)) = op2(h12(e12), h12(e11))) | ~ (h12(op1(e12, e10)) = op2(h12(e12), h12(e10))) | ~ (h12(op1(e11, e15)) = op2(h12(e11), h12(e15))) | ~ (h12(op1(e11, e14)) = op2(h12(e11), h12(e14))) | ~ (h12(op1(e11, e13)) = op2(h12(e11), h12(e13))) | ~ (h12(op1(e11, e12)) = op2(h12(e11), h12(e12))) | ~ (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) | ~ (h12(op1(e11, e10)) = op2(h12(e11), h12(e10))) | ~ (h12(op1(e10, e15)) = op2(h12(e10), h12(e15))) | ~ (h12(op1(e10, e14)) = op2(h12(e10), h12(e14))) | ~ (h12(op1(e10, e13)) = op2(h12(e10), h12(e13))) | ~ (h12(op1(e10, e12)) = op2(h12(e10), h12(e12))) | ~ (h12(op1(e10, e11)) = op2(h12(e10), h12(e11))) | ~ (h12(op1(e10, e10)) = op2(h12(e10), h12(e10)))) & ((~ (e25 = h11(e15)) & ~ (e25 = h11(e14)) & ~ (e25 = h11(e13)) & ~ (e25 = h11(e12)) & ~ (e25 = h11(e11)) & ~ (e25 = h11(e10))) | (~ (e24 = h11(e15)) & ~ (e24 = h11(e14)) & ~ (e24 = h11(e13)) & ~ (e24 = h11(e12)) & ~ (e24 = h11(e11)) & ~ (e24 = h11(e10))) | (~ (e23 = h11(e15)) & ~ (e23 = h11(e14)) & ~ (e23 = h11(e13)) & ~ (e23 = h11(e12)) & ~ (e23 = h11(e11)) & ~ (e23 = h11(e10))) | (~ (e22 = h11(e15)) & ~ (e22 = h11(e14)) & ~ (e22 = h11(e13)) & ~ (e22 = h11(e12)) & ~ (e22 = h11(e11)) & ~ (e22 = h11(e10))) | (~ (e21 = h11(e15)) & ~ (e21 = h11(e14)) & ~ (e21 = h11(e13)) & ~ (e21 = h11(e12)) & ~ (e21 = h11(e11)) & ~ (e21 = h11(e10))) | (~ (e20 = h11(e15)) & ~ (e20 = h11(e14)) & ~ (e20 = h11(e13)) & ~ (e20 = h11(e12)) & ~ (e20 = h11(e11)) & ~ (e20 = h11(e10))) | ~ (h11(op1(e15, e15)) = op2(h11(e15), h11(e15))) | ~ (h11(op1(e15, e14)) = op2(h11(e15), h11(e14))) | ~ (h11(op1(e15, e13)) = op2(h11(e15), h11(e13))) | ~ (h11(op1(e15, e12)) = op2(h11(e15), h11(e12))) | ~ (h11(op1(e15, e11)) = op2(h11(e15), h11(e11))) | ~ (h11(op1(e15, e10)) = op2(h11(e15), h11(e10))) | ~ (h11(op1(e14, e15)) = op2(h11(e14), h11(e15))) | ~ (h11(op1(e14, e14)) = op2(h11(e14), h11(e14))) | ~ (h11(op1(e14, e13)) = op2(h11(e14), h11(e13))) | ~ (h11(op1(e14, e12)) = op2(h11(e14), h11(e12))) | ~ (h11(op1(e14, e11)) = op2(h11(e14), h11(e11))) | ~ (h11(op1(e14, e10)) = op2(h11(e14), h11(e10))) | ~ (h11(op1(e13, e15)) = op2(h11(e13), h11(e15))) | ~ (h11(op1(e13, e14)) = op2(h11(e13), h11(e14))) | ~ (h11(op1(e13, e13)) = op2(h11(e13), h11(e13))) | ~ (h11(op1(e13, e12)) = op2(h11(e13), h11(e12))) | ~ (h11(op1(e13, e11)) = op2(h11(e13), h11(e11))) | ~ (h11(op1(e13, e10)) = op2(h11(e13), h11(e10))) | ~ (h11(op1(e12, e15)) = op2(h11(e12), h11(e15))) | ~ (h11(op1(e12, e14)) = op2(h11(e12), h11(e14))) | ~ (h11(op1(e12, e13)) = op2(h11(e12), h11(e13))) | ~ (h11(op1(e12, e12)) = op2(h11(e12), h11(e12))) | ~ (h11(op1(e12, e11)) = op2(h11(e12), h11(e11))) | ~ (h11(op1(e12, e10)) = op2(h11(e12), h11(e10))) | ~ (h11(op1(e11, e15)) = op2(h11(e11), h11(e15))) | ~ (h11(op1(e11, e14)) = op2(h11(e11), h11(e14))) | ~ (h11(op1(e11, e13)) = op2(h11(e11), h11(e13))) | ~ (h11(op1(e11, e12)) = op2(h11(e11), h11(e12))) | ~ (h11(op1(e11, e11)) = op2(h11(e11), h11(e11))) | ~ (h11(op1(e11, e10)) = op2(h11(e11), h11(e10))) | ~ (h11(op1(e10, e15)) = op2(h11(e10), h11(e15))) | ~ (h11(op1(e10, e14)) = op2(h11(e10), h11(e14))) | ~ (h11(op1(e10, e13)) = op2(h11(e10), h11(e13))) | ~ (h11(op1(e10, e12)) = op2(h11(e10), h11(e12))) | ~ (h11(op1(e10, e11)) = op2(h11(e10), h11(e11))) | ~ (h11(op1(e10, e10)) = op2(h11(e10), h11(e10)))) & ((~ (e25 = h10(e15)) & ~ (e25 = h10(e14)) & ~ (e25 = h10(e13)) & ~ (e25 = h10(e12)) & ~ (e25 = h10(e11)) & ~ (e25 = h10(e10))) | (~ (e24 = h10(e15)) & ~ (e24 = h10(e14)) & ~ (e24 = h10(e13)) & ~ (e24 = h10(e12)) & ~ (e24 = h10(e11)) & ~ (e24 = h10(e10))) | (~ (e23 = h10(e15)) & ~ (e23 = h10(e14)) & ~ (e23 = h10(e13)) & ~ (e23 = h10(e12)) & ~ (e23 = h10(e11)) & ~ (e23 = h10(e10))) | (~ (e22 = h10(e15)) & ~ (e22 = h10(e14)) & ~ (e22 = h10(e13)) & ~ (e22 = h10(e12)) & ~ (e22 = h10(e11)) & ~ (e22 = h10(e10))) | (~ (e21 = h10(e15)) & ~ (e21 = h10(e14)) & ~ (e21 = h10(e13)) & ~ (e21 = h10(e12)) & ~ (e21 = h10(e11)) & ~ (e21 = h10(e10))) | (~ (e20 = h10(e15)) & ~ (e20 = h10(e14)) & ~ (e20 = h10(e13)) & ~ (e20 = h10(e12)) & ~ (e20 = h10(e11)) & ~ (e20 = h10(e10))) | ~ (h10(op1(e15, e15)) = op2(h10(e15), h10(e15))) | ~ (h10(op1(e15, e14)) = op2(h10(e15), h10(e14))) | ~ (h10(op1(e15, e13)) = op2(h10(e15), h10(e13))) | ~ (h10(op1(e15, e12)) = op2(h10(e15), h10(e12))) | ~ (h10(op1(e15, e11)) = op2(h10(e15), h10(e11))) | ~ (h10(op1(e15, e10)) = op2(h10(e15), h10(e10))) | ~ (h10(op1(e14, e15)) = op2(h10(e14), h10(e15))) | ~ (h10(op1(e14, e14)) = op2(h10(e14), h10(e14))) | ~ (h10(op1(e14, e13)) = op2(h10(e14), h10(e13))) | ~ (h10(op1(e14, e12)) = op2(h10(e14), h10(e12))) | ~ (h10(op1(e14, e11)) = op2(h10(e14), h10(e11))) | ~ (h10(op1(e14, e10)) = op2(h10(e14), h10(e10))) | ~ (h10(op1(e13, e15)) = op2(h10(e13), h10(e15))) | ~ (h10(op1(e13, e14)) = op2(h10(e13), h10(e14))) | ~ (h10(op1(e13, e13)) = op2(h10(e13), h10(e13))) | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) | ~ (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) | ~ (h10(op1(e12, e15)) = op2(h10(e12), h10(e15))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), h10(e14))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) | ~ (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) | ~ (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) | ~ (h10(op1(e11, e15)) = op2(h10(e11), h10(e15))) | ~ (h10(op1(e11, e14)) = op2(h10(e11), h10(e14))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e15)) = op2(h10(e10), h10(e15))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))) & ((~ (e25 = h9(e15)) & ~ (e25 = h9(e14)) & ~ (e25 = h9(e13)) & ~ (e25 = h9(e12)) & ~ (e25 = h9(e11)) & ~ (e25 = h9(e10))) | (~ (e24 = h9(e15)) & ~ (e24 = h9(e14)) & ~ (e24 = h9(e13)) & ~ (e24 = h9(e12)) & ~ (e24 = h9(e11)) & ~ (e24 = h9(e10))) | (~ (e23 = h9(e15)) & ~ (e23 = h9(e14)) & ~ (e23 = h9(e13)) & ~ (e23 = h9(e12)) & ~ (e23 = h9(e11)) & ~ (e23 = h9(e10))) | (~ (e22 = h9(e15)) & ~ (e22 = h9(e14)) & ~ (e22 = h9(e13)) & ~ (e22 = h9(e12)) & ~ (e22 = h9(e11)) & ~ (e22 = h9(e10))) | (~ (e21 = h9(e15)) & ~ (e21 = h9(e14)) & ~ (e21 = h9(e13)) & ~ (e21 = h9(e12)) & ~ (e21 = h9(e11)) & ~ (e21 = h9(e10))) | (~ (e20 = h9(e15)) & ~ (e20 = h9(e14)) & ~ (e20 = h9(e13)) & ~ (e20 = h9(e12)) & ~ (e20 = h9(e11)) & ~ (e20 = h9(e10))) | ~ (h9(op1(e15, e15)) = op2(h9(e15), h9(e15))) | ~ (h9(op1(e15, e14)) = op2(h9(e15), h9(e14))) | ~ (h9(op1(e15, e13)) = op2(h9(e15), h9(e13))) | ~ (h9(op1(e15, e12)) = op2(h9(e15), h9(e12))) | ~ (h9(op1(e15, e11)) = op2(h9(e15), h9(e11))) | ~ (h9(op1(e15, e10)) = op2(h9(e15), h9(e10))) | ~ (h9(op1(e14, e15)) = op2(h9(e14), h9(e15))) | ~ (h9(op1(e14, e14)) = op2(h9(e14), h9(e14))) | ~ (h9(op1(e14, e13)) = op2(h9(e14), h9(e13))) | ~ (h9(op1(e14, e12)) = op2(h9(e14), h9(e12))) | ~ (h9(op1(e14, e11)) = op2(h9(e14), h9(e11))) | ~ (h9(op1(e14, e10)) = op2(h9(e14), h9(e10))) | ~ (h9(op1(e13, e15)) = op2(h9(e13), h9(e15))) | ~ (h9(op1(e13, e14)) = op2(h9(e13), h9(e14))) | ~ (h9(op1(e13, e13)) = op2(h9(e13), h9(e13))) | ~ (h9(op1(e13, e12)) = op2(h9(e13), h9(e12))) | ~ (h9(op1(e13, e11)) = op2(h9(e13), h9(e11))) | ~ (h9(op1(e13, e10)) = op2(h9(e13), h9(e10))) | ~ (h9(op1(e12, e15)) = op2(h9(e12), h9(e15))) | ~ (h9(op1(e12, e14)) = op2(h9(e12), h9(e14))) | ~ (h9(op1(e12, e13)) = op2(h9(e12), h9(e13))) | ~ (h9(op1(e12, e12)) = op2(h9(e12), h9(e12))) | ~ (h9(op1(e12, e11)) = op2(h9(e12), h9(e11))) | ~ (h9(op1(e12, e10)) = op2(h9(e12), h9(e10))) | ~ (h9(op1(e11, e15)) = op2(h9(e11), h9(e15))) | ~ (h9(op1(e11, e14)) = op2(h9(e11), h9(e14))) | ~ (h9(op1(e11, e13)) = op2(h9(e11), h9(e13))) | ~ (h9(op1(e11, e12)) = op2(h9(e11), h9(e12))) | ~ (h9(op1(e11, e11)) = op2(h9(e11), h9(e11))) | ~ (h9(op1(e11, e10)) = op2(h9(e11), h9(e10))) | ~ (h9(op1(e10, e15)) = op2(h9(e10), h9(e15))) | ~ (h9(op1(e10, e14)) = op2(h9(e10), h9(e14))) | ~ (h9(op1(e10, e13)) = op2(h9(e10), h9(e13))) | ~ (h9(op1(e10, e12)) = op2(h9(e10), h9(e12))) | ~ (h9(op1(e10, e11)) = op2(h9(e10), h9(e11))) | ~ (h9(op1(e10, e10)) = op2(h9(e10), h9(e10)))) & ((~ (e25 = h8(e15)) & ~ (e25 = h8(e14)) & ~ (e25 = h8(e13)) & ~ (e25 = h8(e12)) & ~ (e25 = h8(e11)) & ~ (e25 = h8(e10))) | (~ (e24 = h8(e15)) & ~ (e24 = h8(e14)) & ~ (e24 = h8(e13)) & ~ (e24 = h8(e12)) & ~ (e24 = h8(e11)) & ~ (e24 = h8(e10))) | (~ (e23 = h8(e15)) & ~ (e23 = h8(e14)) & ~ (e23 = h8(e13)) & ~ (e23 = h8(e12)) & ~ (e23 = h8(e11)) & ~ (e23 = h8(e10))) | (~ (e22 = h8(e15)) & ~ (e22 = h8(e14)) & ~ (e22 = h8(e13)) & ~ (e22 = h8(e12)) & ~ (e22 = h8(e11)) & ~ (e22 = h8(e10))) | (~ (e21 = h8(e15)) & ~ (e21 = h8(e14)) & ~ (e21 = h8(e13)) & ~ (e21 = h8(e12)) & ~ (e21 = h8(e11)) & ~ (e21 = h8(e10))) | (~ (e20 = h8(e15)) & ~ (e20 = h8(e14)) & ~ (e20 = h8(e13)) & ~ (e20 = h8(e12)) & ~ (e20 = h8(e11)) & ~ (e20 = h8(e10))) | ~ (h8(op1(e15, e15)) = op2(h8(e15), h8(e15))) | ~ (h8(op1(e15, e14)) = op2(h8(e15), h8(e14))) | ~ (h8(op1(e15, e13)) = op2(h8(e15), h8(e13))) | ~ (h8(op1(e15, e12)) = op2(h8(e15), h8(e12))) | ~ (h8(op1(e15, e11)) = op2(h8(e15), h8(e11))) | ~ (h8(op1(e15, e10)) = op2(h8(e15), h8(e10))) | ~ (h8(op1(e14, e15)) = op2(h8(e14), h8(e15))) | ~ (h8(op1(e14, e14)) = op2(h8(e14), h8(e14))) | ~ (h8(op1(e14, e13)) = op2(h8(e14), h8(e13))) | ~ (h8(op1(e14, e12)) = op2(h8(e14), h8(e12))) | ~ (h8(op1(e14, e11)) = op2(h8(e14), h8(e11))) | ~ (h8(op1(e14, e10)) = op2(h8(e14), h8(e10))) | ~ (h8(op1(e13, e15)) = op2(h8(e13), h8(e15))) | ~ (h8(op1(e13, e14)) = op2(h8(e13), h8(e14))) | ~ (h8(op1(e13, e13)) = op2(h8(e13), h8(e13))) | ~ (h8(op1(e13, e12)) = op2(h8(e13), h8(e12))) | ~ (h8(op1(e13, e11)) = op2(h8(e13), h8(e11))) | ~ (h8(op1(e13, e10)) = op2(h8(e13), h8(e10))) | ~ (h8(op1(e12, e15)) = op2(h8(e12), h8(e15))) | ~ (h8(op1(e12, e14)) = op2(h8(e12), h8(e14))) | ~ (h8(op1(e12, e13)) = op2(h8(e12), h8(e13))) | ~ (h8(op1(e12, e12)) = op2(h8(e12), h8(e12))) | ~ (h8(op1(e12, e11)) = op2(h8(e12), h8(e11))) | ~ (h8(op1(e12, e10)) = op2(h8(e12), h8(e10))) | ~ (h8(op1(e11, e15)) = op2(h8(e11), h8(e15))) | ~ (h8(op1(e11, e14)) = op2(h8(e11), h8(e14))) | ~ (h8(op1(e11, e13)) = op2(h8(e11), h8(e13))) | ~ (h8(op1(e11, e12)) = op2(h8(e11), h8(e12))) | ~ (h8(op1(e11, e11)) = op2(h8(e11), h8(e11))) | ~ (h8(op1(e11, e10)) = op2(h8(e11), h8(e10))) | ~ (h8(op1(e10, e15)) = op2(h8(e10), h8(e15))) | ~ (h8(op1(e10, e14)) = op2(h8(e10), h8(e14))) | ~ (h8(op1(e10, e13)) = op2(h8(e10), h8(e13))) | ~ (h8(op1(e10, e12)) = op2(h8(e10), h8(e12))) | ~ (h8(op1(e10, e11)) = op2(h8(e10), h8(e11))) | ~ (h8(op1(e10, e10)) = op2(h8(e10), h8(e10)))) & ((~ (e25 = h7(e15)) & ~ (e25 = h7(e14)) & ~ (e25 = h7(e13)) & ~ (e25 = h7(e12)) & ~ (e25 = h7(e11)) & ~ (e25 = h7(e10))) | (~ (e24 = h7(e15)) & ~ (e24 = h7(e14)) & ~ (e24 = h7(e13)) & ~ (e24 = h7(e12)) & ~ (e24 = h7(e11)) & ~ (e24 = h7(e10))) | (~ (e23 = h7(e15)) & ~ (e23 = h7(e14)) & ~ (e23 = h7(e13)) & ~ (e23 = h7(e12)) & ~ (e23 = h7(e11)) & ~ (e23 = h7(e10))) | (~ (e22 = h7(e15)) & ~ (e22 = h7(e14)) & ~ (e22 = h7(e13)) & ~ (e22 = h7(e12)) & ~ (e22 = h7(e11)) & ~ (e22 = h7(e10))) | (~ (e21 = h7(e15)) & ~ (e21 = h7(e14)) & ~ (e21 = h7(e13)) & ~ (e21 = h7(e12)) & ~ (e21 = h7(e11)) & ~ (e21 = h7(e10))) | (~ (e20 = h7(e15)) & ~ (e20 = h7(e14)) & ~ (e20 = h7(e13)) & ~ (e20 = h7(e12)) & ~ (e20 = h7(e11)) & ~ (e20 = h7(e10))) | ~ (h7(op1(e15, e15)) = op2(h7(e15), h7(e15))) | ~ (h7(op1(e15, e14)) = op2(h7(e15), h7(e14))) | ~ (h7(op1(e15, e13)) = op2(h7(e15), h7(e13))) | ~ (h7(op1(e15, e12)) = op2(h7(e15), h7(e12))) | ~ (h7(op1(e15, e11)) = op2(h7(e15), h7(e11))) | ~ (h7(op1(e15, e10)) = op2(h7(e15), h7(e10))) | ~ (h7(op1(e14, e15)) = op2(h7(e14), h7(e15))) | ~ (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) | ~ (h7(op1(e14, e13)) = op2(h7(e14), h7(e13))) | ~ (h7(op1(e14, e12)) = op2(h7(e14), h7(e12))) | ~ (h7(op1(e14, e11)) = op2(h7(e14), h7(e11))) | ~ (h7(op1(e14, e10)) = op2(h7(e14), h7(e10))) | ~ (h7(op1(e13, e15)) = op2(h7(e13), h7(e15))) | ~ (h7(op1(e13, e14)) = op2(h7(e13), h7(e14))) | ~ (h7(op1(e13, e13)) = op2(h7(e13), h7(e13))) | ~ (h7(op1(e13, e12)) = op2(h7(e13), h7(e12))) | ~ (h7(op1(e13, e11)) = op2(h7(e13), h7(e11))) | ~ (h7(op1(e13, e10)) = op2(h7(e13), h7(e10))) | ~ (h7(op1(e12, e15)) = op2(h7(e12), h7(e15))) | ~ (h7(op1(e12, e14)) = op2(h7(e12), h7(e14))) | ~ (h7(op1(e12, e13)) = op2(h7(e12), h7(e13))) | ~ (h7(op1(e12, e12)) = op2(h7(e12), h7(e12))) | ~ (h7(op1(e12, e11)) = op2(h7(e12), h7(e11))) | ~ (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) | ~ (h7(op1(e11, e15)) = op2(h7(e11), h7(e15))) | ~ (h7(op1(e11, e14)) = op2(h7(e11), h7(e14))) | ~ (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) | ~ (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) | ~ (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) | ~ (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) | ~ (h7(op1(e10, e15)) = op2(h7(e10), h7(e15))) | ~ (h7(op1(e10, e14)) = op2(h7(e10), h7(e14))) | ~ (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) | ~ (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) | ~ (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))) & ((~ (e25 = h6(e15)) & ~ (e25 = h6(e14)) & ~ (e25 = h6(e13)) & ~ (e25 = h6(e12)) & ~ (e25 = h6(e11)) & ~ (e25 = h6(e10))) | (~ (e24 = h6(e15)) & ~ (e24 = h6(e14)) & ~ (e24 = h6(e13)) & ~ (e24 = h6(e12)) & ~ (e24 = h6(e11)) & ~ (e24 = h6(e10))) | (~ (e23 = h6(e15)) & ~ (e23 = h6(e14)) & ~ (e23 = h6(e13)) & ~ (e23 = h6(e12)) & ~ (e23 = h6(e11)) & ~ (e23 = h6(e10))) | (~ (e22 = h6(e15)) & ~ (e22 = h6(e14)) & ~ (e22 = h6(e13)) & ~ (e22 = h6(e12)) & ~ (e22 = h6(e11)) & ~ (e22 = h6(e10))) | (~ (e21 = h6(e15)) & ~ (e21 = h6(e14)) & ~ (e21 = h6(e13)) & ~ (e21 = h6(e12)) & ~ (e21 = h6(e11)) & ~ (e21 = h6(e10))) | (~ (e20 = h6(e15)) & ~ (e20 = h6(e14)) & ~ (e20 = h6(e13)) & ~ (e20 = h6(e12)) & ~ (e20 = h6(e11)) & ~ (e20 = h6(e10))) | ~ (h6(op1(e15, e15)) = op2(h6(e15), h6(e15))) | ~ (h6(op1(e15, e14)) = op2(h6(e15), h6(e14))) | ~ (h6(op1(e15, e13)) = op2(h6(e15), h6(e13))) | ~ (h6(op1(e15, e12)) = op2(h6(e15), h6(e12))) | ~ (h6(op1(e15, e11)) = op2(h6(e15), h6(e11))) | ~ (h6(op1(e15, e10)) = op2(h6(e15), h6(e10))) | ~ (h6(op1(e14, e15)) = op2(h6(e14), h6(e15))) | ~ (h6(op1(e14, e14)) = op2(h6(e14), h6(e14))) | ~ (h6(op1(e14, e13)) = op2(h6(e14), h6(e13))) | ~ (h6(op1(e14, e12)) = op2(h6(e14), h6(e12))) | ~ (h6(op1(e14, e11)) = op2(h6(e14), h6(e11))) | ~ (h6(op1(e14, e10)) = op2(h6(e14), h6(e10))) | ~ (h6(op1(e13, e15)) = op2(h6(e13), h6(e15))) | ~ (h6(op1(e13, e14)) = op2(h6(e13), h6(e14))) | ~ (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) | ~ (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e15)) = op2(h6(e12), h6(e15))) | ~ (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))) & ((~ (e25 = h5(e15)) & ~ (e25 = h5(e14)) & ~ (e25 = h5(e13)) & ~ (e25 = h5(e12)) & ~ (e25 = h5(e11)) & ~ (e25 = h5(e10))) | (~ (e24 = h5(e15)) & ~ (e24 = h5(e14)) & ~ (e24 = h5(e13)) & ~ (e24 = h5(e12)) & ~ (e24 = h5(e11)) & ~ (e24 = h5(e10))) | (~ (e23 = h5(e15)) & ~ (e23 = h5(e14)) & ~ (e23 = h5(e13)) & ~ (e23 = h5(e12)) & ~ (e23 = h5(e11)) & ~ (e23 = h5(e10))) | (~ (e22 = h5(e15)) & ~ (e22 = h5(e14)) & ~ (e22 = h5(e13)) & ~ (e22 = h5(e12)) & ~ (e22 = h5(e11)) & ~ (e22 = h5(e10))) | (~ (e21 = h5(e15)) & ~ (e21 = h5(e14)) & ~ (e21 = h5(e13)) & ~ (e21 = h5(e12)) & ~ (e21 = h5(e11)) & ~ (e21 = h5(e10))) | (~ (e20 = h5(e15)) & ~ (e20 = h5(e14)) & ~ (e20 = h5(e13)) & ~ (e20 = h5(e12)) & ~ (e20 = h5(e11)) & ~ (e20 = h5(e10))) | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e15, e14)) = op2(h5(e15), h5(e14))) | ~ (h5(op1(e15, e13)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e15, e12)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e15, e11)) = op2(h5(e15), h5(e11))) | ~ (h5(op1(e15, e10)) = op2(h5(e15), h5(e10))) | ~ (h5(op1(e14, e15)) = op2(h5(e14), h5(e15))) | ~ (h5(op1(e14, e14)) = op2(h5(e14), h5(e14))) | ~ (h5(op1(e14, e13)) = op2(h5(e14), h5(e13))) | ~ (h5(op1(e14, e12)) = op2(h5(e14), h5(e12))) | ~ (h5(op1(e14, e11)) = op2(h5(e14), h5(e11))) | ~ (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e15)) = op2(h5(e11), h5(e15))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))) & ((~ (e25 = h4(e15)) & ~ (e25 = h4(e14)) & ~ (e25 = h4(e13)) & ~ (e25 = h4(e12)) & ~ (e25 = h4(e11)) & ~ (e25 = h4(e10))) | (~ (e24 = h4(e15)) & ~ (e24 = h4(e14)) & ~ (e24 = h4(e13)) & ~ (e24 = h4(e12)) & ~ (e24 = h4(e11)) & ~ (e24 = h4(e10))) | (~ (e23 = h4(e15)) & ~ (e23 = h4(e14)) & ~ (e23 = h4(e13)) & ~ (e23 = h4(e12)) & ~ (e23 = h4(e11)) & ~ (e23 = h4(e10))) | (~ (e22 = h4(e15)) & ~ (e22 = h4(e14)) & ~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10))) | (~ (e21 = h4(e15)) & ~ (e21 = h4(e14)) & ~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10))) | (~ (e20 = h4(e15)) & ~ (e20 = h4(e14)) & ~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10))) | ~ (h4(op1(e15, e15)) = op2(h4(e15), h4(e15))) | ~ (h4(op1(e15, e14)) = op2(h4(e15), h4(e14))) | ~ (h4(op1(e15, e13)) = op2(h4(e15), h4(e13))) | ~ (h4(op1(e15, e12)) = op2(h4(e15), h4(e12))) | ~ (h4(op1(e15, e11)) = op2(h4(e15), h4(e11))) | ~ (h4(op1(e15, e10)) = op2(h4(e15), h4(e10))) | ~ (h4(op1(e14, e15)) = op2(h4(e14), h4(e15))) | ~ (h4(op1(e14, e14)) = op2(h4(e14), h4(e14))) | ~ (h4(op1(e14, e13)) = op2(h4(e14), h4(e13))) | ~ (h4(op1(e14, e12)) = op2(h4(e14), h4(e12))) | ~ (h4(op1(e14, e11)) = op2(h4(e14), h4(e11))) | ~ (h4(op1(e14, e10)) = op2(h4(e14), h4(e10))) | ~ (h4(op1(e13, e15)) = op2(h4(e13), h4(e15))) | ~ (h4(op1(e13, e14)) = op2(h4(e13), h4(e14))) | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e15)) = op2(h4(e12), h4(e15))) | ~ (h4(op1(e12, e14)) = op2(h4(e12), h4(e14))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e15)) = op2(h4(e11), h4(e15))) | ~ (h4(op1(e11, e14)) = op2(h4(e11), h4(e14))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e15)) = op2(h4(e10), h4(e15))) | ~ (h4(op1(e10, e14)) = op2(h4(e10), h4(e14))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) & ((~ (e25 = h3(e15)) & ~ (e25 = h3(e14)) & ~ (e25 = h3(e13)) & ~ (e25 = h3(e12)) & ~ (e25 = h3(e11)) & ~ (e25 = h3(e10))) | (~ (e24 = h3(e15)) & ~ (e24 = h3(e14)) & ~ (e24 = h3(e13)) & ~ (e24 = h3(e12)) & ~ (e24 = h3(e11)) & ~ (e24 = h3(e10))) | (~ (e23 = h3(e15)) & ~ (e23 = h3(e14)) & ~ (e23 = h3(e13)) & ~ (e23 = h3(e12)) & ~ (e23 = h3(e11)) & ~ (e23 = h3(e10))) | (~ (e22 = h3(e15)) & ~ (e22 = h3(e14)) & ~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10))) | (~ (e21 = h3(e15)) & ~ (e21 = h3(e14)) & ~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10))) | (~ (e20 = h3(e15)) & ~ (e20 = h3(e14)) & ~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10))) | ~ (h3(op1(e15, e15)) = op2(h3(e15), h3(e15))) | ~ (h3(op1(e15, e14)) = op2(h3(e15), h3(e14))) | ~ (h3(op1(e15, e13)) = op2(h3(e15), h3(e13))) | ~ (h3(op1(e15, e12)) = op2(h3(e15), h3(e12))) | ~ (h3(op1(e15, e11)) = op2(h3(e15), h3(e11))) | ~ (h3(op1(e15, e10)) = op2(h3(e15), h3(e10))) | ~ (h3(op1(e14, e15)) = op2(h3(e14), h3(e15))) | ~ (h3(op1(e14, e14)) = op2(h3(e14), h3(e14))) | ~ (h3(op1(e14, e13)) = op2(h3(e14), h3(e13))) | ~ (h3(op1(e14, e12)) = op2(h3(e14), h3(e12))) | ~ (h3(op1(e14, e11)) = op2(h3(e14), h3(e11))) | ~ (h3(op1(e14, e10)) = op2(h3(e14), h3(e10))) | ~ (h3(op1(e13, e15)) = op2(h3(e13), h3(e15))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), h3(e14))) | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e15)) = op2(h3(e12), h3(e15))) | ~ (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e15)) = op2(h3(e11), h3(e15))) | ~ (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e15)) = op2(h3(e10), h3(e15))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) & ((~ (e25 = h2(e15)) & ~ (e25 = h2(e14)) & ~ (e25 = h2(e13)) & ~ (e25 = h2(e12)) & ~ (e25 = h2(e11)) & ~ (e25 = h2(e10))) | (~ (e24 = h2(e15)) & ~ (e24 = h2(e14)) & ~ (e24 = h2(e13)) & ~ (e24 = h2(e12)) & ~ (e24 = h2(e11)) & ~ (e24 = h2(e10))) | (~ (e23 = h2(e15)) & ~ (e23 = h2(e14)) & ~ (e23 = h2(e13)) & ~ (e23 = h2(e12)) & ~ (e23 = h2(e11)) & ~ (e23 = h2(e10))) | (~ (e22 = h2(e15)) & ~ (e22 = h2(e14)) & ~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10))) | (~ (e21 = h2(e15)) & ~ (e21 = h2(e14)) & ~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10))) | (~ (e20 = h2(e15)) & ~ (e20 = h2(e14)) & ~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10))) | ~ (h2(op1(e15, e15)) = op2(h2(e15), h2(e15))) | ~ (h2(op1(e15, e14)) = op2(h2(e15), h2(e14))) | ~ (h2(op1(e15, e13)) = op2(h2(e15), h2(e13))) | ~ (h2(op1(e15, e12)) = op2(h2(e15), h2(e12))) | ~ (h2(op1(e15, e11)) = op2(h2(e15), h2(e11))) | ~ (h2(op1(e15, e10)) = op2(h2(e15), h2(e10))) | ~ (h2(op1(e14, e15)) = op2(h2(e14), h2(e15))) | ~ (h2(op1(e14, e14)) = op2(h2(e14), h2(e14))) | ~ (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) | ~ (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) | ~ (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) | ~ (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) | ~ (h2(op1(e13, e15)) = op2(h2(e13), h2(e15))) | ~ (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e15)) = op2(h2(e12), h2(e15))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e15)) = op2(h2(e11), h2(e15))) | ~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e15)) = op2(h2(e10), h2(e15))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) & ((~ (e25 = h1(e15)) & ~ (e25 = h1(e14)) & ~ (e25 = h1(e13)) & ~ (e25 = h1(e12)) & ~ (e25 = h1(e11)) & ~ (e25 = h1(e10))) | (~ (e24 = h1(e15)) & ~ (e24 = h1(e14)) & ~ (e24 = h1(e13)) & ~ (e24 = h1(e12)) & ~ (e24 = h1(e11)) & ~ (e24 = h1(e10))) | (~ (e23 = h1(e15)) & ~ (e23 = h1(e14)) & ~ (e23 = h1(e13)) & ~ (e23 = h1(e12)) & ~ (e23 = h1(e11)) & ~ (e23 = h1(e10))) | (~ (e22 = h1(e15)) & ~ (e22 = h1(e14)) & ~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10))) | (~ (e21 = h1(e15)) & ~ (e21 = h1(e14)) & ~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10))) | (~ (e20 = h1(e15)) & ~ (e20 = h1(e14)) & ~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10))) | ~ (h1(op1(e15, e15)) = op2(h1(e15), h1(e15))) | ~ (h1(op1(e15, e14)) = op2(h1(e15), h1(e14))) | ~ (h1(op1(e15, e13)) = op2(h1(e15), h1(e13))) | ~ (h1(op1(e15, e12)) = op2(h1(e15), h1(e12))) | ~ (h1(op1(e15, e11)) = op2(h1(e15), h1(e11))) | ~ (h1(op1(e15, e10)) = op2(h1(e15), h1(e10))) | ~ (h1(op1(e14, e15)) = op2(h1(e14), h1(e15))) | ~ (h1(op1(e14, e14)) = op2(h1(e14), h1(e14))) | ~ (h1(op1(e14, e13)) = op2(h1(e14), h1(e13))) | ~ (h1(op1(e14, e12)) = op2(h1(e14), h1(e12))) | ~ (h1(op1(e14, e11)) = op2(h1(e14), h1(e11))) | ~ (h1(op1(e14, e10)) = op2(h1(e14), h1(e10))) | ~ (h1(op1(e13, e15)) = op2(h1(e13), h1(e15))) | ~ (h1(op1(e13, e14)) = op2(h1(e13), h1(e14))) | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e15)) = op2(h1(e12), h1(e15))) | ~ (h1(op1(e12, e14)) = op2(h1(e12), h1(e14))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e15)) = op2(h1(e11), h1(e15))) | ~ (h1(op1(e11, e14)) = op2(h1(e11), h1(e14))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e15)) = op2(h1(e10), h1(e15))) | ~ (h1(op1(e10, e14)) = op2(h1(e10), h1(e14))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(ennf_transformation, [], [f57])).
fof(f57, plain, ~ ((((e25 = h30(e15)) | (e25 = h30(e14)) | (e25 = h30(e13)) | (e25 = h30(e12)) | (e25 = h30(e11)) | (e25 = h30(e10))) & ((e24 = h30(e15)) | (e24 = h30(e14)) | (e24 = h30(e13)) | (e24 = h30(e12)) | (e24 = h30(e11)) | (e24 = h30(e10))) & ((e23 = h30(e15)) | (e23 = h30(e14)) | (e23 = h30(e13)) | (e23 = h30(e12)) | (e23 = h30(e11)) | (e23 = h30(e10))) & ((e22 = h30(e15)) | (e22 = h30(e14)) | (e22 = h30(e13)) | (e22 = h30(e12)) | (e22 = h30(e11)) | (e22 = h30(e10))) & ((e21 = h30(e15)) | (e21 = h30(e14)) | (e21 = h30(e13)) | (e21 = h30(e12)) | (e21 = h30(e11)) | (e21 = h30(e10))) & ((e20 = h30(e15)) | (e20 = h30(e14)) | (e20 = h30(e13)) | (e20 = h30(e12)) | (e20 = h30(e11)) | (e20 = h30(e10))) & (h30(op1(e15, e15)) = op2(h30(e15), h30(e15))) & (h30(op1(e15, e14)) = op2(h30(e15), h30(e14))) & (h30(op1(e15, e13)) = op2(h30(e15), h30(e13))) & (h30(op1(e15, e12)) = op2(h30(e15), h30(e12))) & (h30(op1(e15, e11)) = op2(h30(e15), h30(e11))) & (h30(op1(e15, e10)) = op2(h30(e15), h30(e10))) & (h30(op1(e14, e15)) = op2(h30(e14), h30(e15))) & (h30(op1(e14, e14)) = op2(h30(e14), h30(e14))) & (h30(op1(e14, e13)) = op2(h30(e14), h30(e13))) & (h30(op1(e14, e12)) = op2(h30(e14), h30(e12))) & (h30(op1(e14, e11)) = op2(h30(e14), h30(e11))) & (h30(op1(e14, e10)) = op2(h30(e14), h30(e10))) & (h30(op1(e13, e15)) = op2(h30(e13), h30(e15))) & (h30(op1(e13, e14)) = op2(h30(e13), h30(e14))) & (h30(op1(e13, e13)) = op2(h30(e13), h30(e13))) & (h30(op1(e13, e12)) = op2(h30(e13), h30(e12))) & (h30(op1(e13, e11)) = op2(h30(e13), h30(e11))) & (h30(op1(e13, e10)) = op2(h30(e13), h30(e10))) & (h30(op1(e12, e15)) = op2(h30(e12), h30(e15))) & (h30(op1(e12, e14)) = op2(h30(e12), h30(e14))) & (h30(op1(e12, e13)) = op2(h30(e12), h30(e13))) & (h30(op1(e12, e12)) = op2(h30(e12), h30(e12))) & (h30(op1(e12, e11)) = op2(h30(e12), h30(e11))) & (h30(op1(e12, e10)) = op2(h30(e12), h30(e10))) & (h30(op1(e11, e15)) = op2(h30(e11), h30(e15))) & (h30(op1(e11, e14)) = op2(h30(e11), h30(e14))) & (h30(op1(e11, e13)) = op2(h30(e11), h30(e13))) & (h30(op1(e11, e12)) = op2(h30(e11), h30(e12))) & (h30(op1(e11, e11)) = op2(h30(e11), h30(e11))) & (h30(op1(e11, e10)) = op2(h30(e11), h30(e10))) & (h30(op1(e10, e15)) = op2(h30(e10), h30(e15))) & (h30(op1(e10, e14)) = op2(h30(e10), h30(e14))) & (h30(op1(e10, e13)) = op2(h30(e10), h30(e13))) & (h30(op1(e10, e12)) = op2(h30(e10), h30(e12))) & (h30(op1(e10, e11)) = op2(h30(e10), h30(e11))) & (h30(op1(e10, e10)) = op2(h30(e10), h30(e10)))) | (((e25 = h29(e15)) | (e25 = h29(e14)) | (e25 = h29(e13)) | (e25 = h29(e12)) | (e25 = h29(e11)) | (e25 = h29(e10))) & ((e24 = h29(e15)) | (e24 = h29(e14)) | (e24 = h29(e13)) | (e24 = h29(e12)) | (e24 = h29(e11)) | (e24 = h29(e10))) & ((e23 = h29(e15)) | (e23 = h29(e14)) | (e23 = h29(e13)) | (e23 = h29(e12)) | (e23 = h29(e11)) | (e23 = h29(e10))) & ((e22 = h29(e15)) | (e22 = h29(e14)) | (e22 = h29(e13)) | (e22 = h29(e12)) | (e22 = h29(e11)) | (e22 = h29(e10))) & ((e21 = h29(e15)) | (e21 = h29(e14)) | (e21 = h29(e13)) | (e21 = h29(e12)) | (e21 = h29(e11)) | (e21 = h29(e10))) & ((e20 = h29(e15)) | (e20 = h29(e14)) | (e20 = h29(e13)) | (e20 = h29(e12)) | (e20 = h29(e11)) | (e20 = h29(e10))) & (h29(op1(e15, e15)) = op2(h29(e15), h29(e15))) & (h29(op1(e15, e14)) = op2(h29(e15), h29(e14))) & (h29(op1(e15, e13)) = op2(h29(e15), h29(e13))) & (h29(op1(e15, e12)) = op2(h29(e15), h29(e12))) & (h29(op1(e15, e11)) = op2(h29(e15), h29(e11))) & (h29(op1(e15, e10)) = op2(h29(e15), h29(e10))) & (h29(op1(e14, e15)) = op2(h29(e14), h29(e15))) & (h29(op1(e14, e14)) = op2(h29(e14), h29(e14))) & (h29(op1(e14, e13)) = op2(h29(e14), h29(e13))) & (h29(op1(e14, e12)) = op2(h29(e14), h29(e12))) & (h29(op1(e14, e11)) = op2(h29(e14), h29(e11))) & (h29(op1(e14, e10)) = op2(h29(e14), h29(e10))) & (h29(op1(e13, e15)) = op2(h29(e13), h29(e15))) & (h29(op1(e13, e14)) = op2(h29(e13), h29(e14))) & (h29(op1(e13, e13)) = op2(h29(e13), h29(e13))) & (h29(op1(e13, e12)) = op2(h29(e13), h29(e12))) & (h29(op1(e13, e11)) = op2(h29(e13), h29(e11))) & (h29(op1(e13, e10)) = op2(h29(e13), h29(e10))) & (h29(op1(e12, e15)) = op2(h29(e12), h29(e15))) & (h29(op1(e12, e14)) = op2(h29(e12), h29(e14))) & (h29(op1(e12, e13)) = op2(h29(e12), h29(e13))) & (h29(op1(e12, e12)) = op2(h29(e12), h29(e12))) & (h29(op1(e12, e11)) = op2(h29(e12), h29(e11))) & (h29(op1(e12, e10)) = op2(h29(e12), h29(e10))) & (h29(op1(e11, e15)) = op2(h29(e11), h29(e15))) & (h29(op1(e11, e14)) = op2(h29(e11), h29(e14))) & (h29(op1(e11, e13)) = op2(h29(e11), h29(e13))) & (h29(op1(e11, e12)) = op2(h29(e11), h29(e12))) & (h29(op1(e11, e11)) = op2(h29(e11), h29(e11))) & (h29(op1(e11, e10)) = op2(h29(e11), h29(e10))) & (h29(op1(e10, e15)) = op2(h29(e10), h29(e15))) & (h29(op1(e10, e14)) = op2(h29(e10), h29(e14))) & (h29(op1(e10, e13)) = op2(h29(e10), h29(e13))) & (h29(op1(e10, e12)) = op2(h29(e10), h29(e12))) & (h29(op1(e10, e11)) = op2(h29(e10), h29(e11))) & (h29(op1(e10, e10)) = op2(h29(e10), h29(e10)))) | (((e25 = h28(e15)) | (e25 = h28(e14)) | (e25 = h28(e13)) | (e25 = h28(e12)) | (e25 = h28(e11)) | (e25 = h28(e10))) & ((e24 = h28(e15)) | (e24 = h28(e14)) | (e24 = h28(e13)) | (e24 = h28(e12)) | (e24 = h28(e11)) | (e24 = h28(e10))) & ((e23 = h28(e15)) | (e23 = h28(e14)) | (e23 = h28(e13)) | (e23 = h28(e12)) | (e23 = h28(e11)) | (e23 = h28(e10))) & ((e22 = h28(e15)) | (e22 = h28(e14)) | (e22 = h28(e13)) | (e22 = h28(e12)) | (e22 = h28(e11)) | (e22 = h28(e10))) & ((e21 = h28(e15)) | (e21 = h28(e14)) | (e21 = h28(e13)) | (e21 = h28(e12)) | (e21 = h28(e11)) | (e21 = h28(e10))) & ((e20 = h28(e15)) | (e20 = h28(e14)) | (e20 = h28(e13)) | (e20 = h28(e12)) | (e20 = h28(e11)) | (e20 = h28(e10))) & (h28(op1(e15, e15)) = op2(h28(e15), h28(e15))) & (h28(op1(e15, e14)) = op2(h28(e15), h28(e14))) & (h28(op1(e15, e13)) = op2(h28(e15), h28(e13))) & (h28(op1(e15, e12)) = op2(h28(e15), h28(e12))) & (h28(op1(e15, e11)) = op2(h28(e15), h28(e11))) & (h28(op1(e15, e10)) = op2(h28(e15), h28(e10))) & (h28(op1(e14, e15)) = op2(h28(e14), h28(e15))) & (h28(op1(e14, e14)) = op2(h28(e14), h28(e14))) & (h28(op1(e14, e13)) = op2(h28(e14), h28(e13))) & (h28(op1(e14, e12)) = op2(h28(e14), h28(e12))) & (h28(op1(e14, e11)) = op2(h28(e14), h28(e11))) & (h28(op1(e14, e10)) = op2(h28(e14), h28(e10))) & (h28(op1(e13, e15)) = op2(h28(e13), h28(e15))) & (h28(op1(e13, e14)) = op2(h28(e13), h28(e14))) & (h28(op1(e13, e13)) = op2(h28(e13), h28(e13))) & (h28(op1(e13, e12)) = op2(h28(e13), h28(e12))) & (h28(op1(e13, e11)) = op2(h28(e13), h28(e11))) & (h28(op1(e13, e10)) = op2(h28(e13), h28(e10))) & (h28(op1(e12, e15)) = op2(h28(e12), h28(e15))) & (h28(op1(e12, e14)) = op2(h28(e12), h28(e14))) & (h28(op1(e12, e13)) = op2(h28(e12), h28(e13))) & (h28(op1(e12, e12)) = op2(h28(e12), h28(e12))) & (h28(op1(e12, e11)) = op2(h28(e12), h28(e11))) & (h28(op1(e12, e10)) = op2(h28(e12), h28(e10))) & (h28(op1(e11, e15)) = op2(h28(e11), h28(e15))) & (h28(op1(e11, e14)) = op2(h28(e11), h28(e14))) & (h28(op1(e11, e13)) = op2(h28(e11), h28(e13))) & (h28(op1(e11, e12)) = op2(h28(e11), h28(e12))) & (h28(op1(e11, e11)) = op2(h28(e11), h28(e11))) & (h28(op1(e11, e10)) = op2(h28(e11), h28(e10))) & (h28(op1(e10, e15)) = op2(h28(e10), h28(e15))) & (h28(op1(e10, e14)) = op2(h28(e10), h28(e14))) & (h28(op1(e10, e13)) = op2(h28(e10), h28(e13))) & (h28(op1(e10, e12)) = op2(h28(e10), h28(e12))) & (h28(op1(e10, e11)) = op2(h28(e10), h28(e11))) & (h28(op1(e10, e10)) = op2(h28(e10), h28(e10)))) | (((e25 = h27(e15)) | (e25 = h27(e14)) | (e25 = h27(e13)) | (e25 = h27(e12)) | (e25 = h27(e11)) | (e25 = h27(e10))) & ((e24 = h27(e15)) | (e24 = h27(e14)) | (e24 = h27(e13)) | (e24 = h27(e12)) | (e24 = h27(e11)) | (e24 = h27(e10))) & ((e23 = h27(e15)) | (e23 = h27(e14)) | (e23 = h27(e13)) | (e23 = h27(e12)) | (e23 = h27(e11)) | (e23 = h27(e10))) & ((e22 = h27(e15)) | (e22 = h27(e14)) | (e22 = h27(e13)) | (e22 = h27(e12)) | (e22 = h27(e11)) | (e22 = h27(e10))) & ((e21 = h27(e15)) | (e21 = h27(e14)) | (e21 = h27(e13)) | (e21 = h27(e12)) | (e21 = h27(e11)) | (e21 = h27(e10))) & ((e20 = h27(e15)) | (e20 = h27(e14)) | (e20 = h27(e13)) | (e20 = h27(e12)) | (e20 = h27(e11)) | (e20 = h27(e10))) & (h27(op1(e15, e15)) = op2(h27(e15), h27(e15))) & (h27(op1(e15, e14)) = op2(h27(e15), h27(e14))) & (h27(op1(e15, e13)) = op2(h27(e15), h27(e13))) & (h27(op1(e15, e12)) = op2(h27(e15), h27(e12))) & (h27(op1(e15, e11)) = op2(h27(e15), h27(e11))) & (h27(op1(e15, e10)) = op2(h27(e15), h27(e10))) & (h27(op1(e14, e15)) = op2(h27(e14), h27(e15))) & (h27(op1(e14, e14)) = op2(h27(e14), h27(e14))) & (h27(op1(e14, e13)) = op2(h27(e14), h27(e13))) & (h27(op1(e14, e12)) = op2(h27(e14), h27(e12))) & (h27(op1(e14, e11)) = op2(h27(e14), h27(e11))) & (h27(op1(e14, e10)) = op2(h27(e14), h27(e10))) & (h27(op1(e13, e15)) = op2(h27(e13), h27(e15))) & (h27(op1(e13, e14)) = op2(h27(e13), h27(e14))) & (h27(op1(e13, e13)) = op2(h27(e13), h27(e13))) & (h27(op1(e13, e12)) = op2(h27(e13), h27(e12))) & (h27(op1(e13, e11)) = op2(h27(e13), h27(e11))) & (h27(op1(e13, e10)) = op2(h27(e13), h27(e10))) & (h27(op1(e12, e15)) = op2(h27(e12), h27(e15))) & (h27(op1(e12, e14)) = op2(h27(e12), h27(e14))) & (h27(op1(e12, e13)) = op2(h27(e12), h27(e13))) & (h27(op1(e12, e12)) = op2(h27(e12), h27(e12))) & (h27(op1(e12, e11)) = op2(h27(e12), h27(e11))) & (h27(op1(e12, e10)) = op2(h27(e12), h27(e10))) & (h27(op1(e11, e15)) = op2(h27(e11), h27(e15))) & (h27(op1(e11, e14)) = op2(h27(e11), h27(e14))) & (h27(op1(e11, e13)) = op2(h27(e11), h27(e13))) & (h27(op1(e11, e12)) = op2(h27(e11), h27(e12))) & (h27(op1(e11, e11)) = op2(h27(e11), h27(e11))) & (h27(op1(e11, e10)) = op2(h27(e11), h27(e10))) & (h27(op1(e10, e15)) = op2(h27(e10), h27(e15))) & (h27(op1(e10, e14)) = op2(h27(e10), h27(e14))) & (h27(op1(e10, e13)) = op2(h27(e10), h27(e13))) & (h27(op1(e10, e12)) = op2(h27(e10), h27(e12))) & (h27(op1(e10, e11)) = op2(h27(e10), h27(e11))) & (h27(op1(e10, e10)) = op2(h27(e10), h27(e10)))) | (((e25 = h26(e15)) | (e25 = h26(e14)) | (e25 = h26(e13)) | (e25 = h26(e12)) | (e25 = h26(e11)) | (e25 = h26(e10))) & ((e24 = h26(e15)) | (e24 = h26(e14)) | (e24 = h26(e13)) | (e24 = h26(e12)) | (e24 = h26(e11)) | (e24 = h26(e10))) & ((e23 = h26(e15)) | (e23 = h26(e14)) | (e23 = h26(e13)) | (e23 = h26(e12)) | (e23 = h26(e11)) | (e23 = h26(e10))) & ((e22 = h26(e15)) | (e22 = h26(e14)) | (e22 = h26(e13)) | (e22 = h26(e12)) | (e22 = h26(e11)) | (e22 = h26(e10))) & ((e21 = h26(e15)) | (e21 = h26(e14)) | (e21 = h26(e13)) | (e21 = h26(e12)) | (e21 = h26(e11)) | (e21 = h26(e10))) & ((e20 = h26(e15)) | (e20 = h26(e14)) | (e20 = h26(e13)) | (e20 = h26(e12)) | (e20 = h26(e11)) | (e20 = h26(e10))) & (h26(op1(e15, e15)) = op2(h26(e15), h26(e15))) & (h26(op1(e15, e14)) = op2(h26(e15), h26(e14))) & (h26(op1(e15, e13)) = op2(h26(e15), h26(e13))) & (h26(op1(e15, e12)) = op2(h26(e15), h26(e12))) & (h26(op1(e15, e11)) = op2(h26(e15), h26(e11))) & (h26(op1(e15, e10)) = op2(h26(e15), h26(e10))) & (h26(op1(e14, e15)) = op2(h26(e14), h26(e15))) & (h26(op1(e14, e14)) = op2(h26(e14), h26(e14))) & (h26(op1(e14, e13)) = op2(h26(e14), h26(e13))) & (h26(op1(e14, e12)) = op2(h26(e14), h26(e12))) & (h26(op1(e14, e11)) = op2(h26(e14), h26(e11))) & (h26(op1(e14, e10)) = op2(h26(e14), h26(e10))) & (h26(op1(e13, e15)) = op2(h26(e13), h26(e15))) & (h26(op1(e13, e14)) = op2(h26(e13), h26(e14))) & (h26(op1(e13, e13)) = op2(h26(e13), h26(e13))) & (h26(op1(e13, e12)) = op2(h26(e13), h26(e12))) & (h26(op1(e13, e11)) = op2(h26(e13), h26(e11))) & (h26(op1(e13, e10)) = op2(h26(e13), h26(e10))) & (h26(op1(e12, e15)) = op2(h26(e12), h26(e15))) & (h26(op1(e12, e14)) = op2(h26(e12), h26(e14))) & (h26(op1(e12, e13)) = op2(h26(e12), h26(e13))) & (h26(op1(e12, e12)) = op2(h26(e12), h26(e12))) & (h26(op1(e12, e11)) = op2(h26(e12), h26(e11))) & (h26(op1(e12, e10)) = op2(h26(e12), h26(e10))) & (h26(op1(e11, e15)) = op2(h26(e11), h26(e15))) & (h26(op1(e11, e14)) = op2(h26(e11), h26(e14))) & (h26(op1(e11, e13)) = op2(h26(e11), h26(e13))) & (h26(op1(e11, e12)) = op2(h26(e11), h26(e12))) & (h26(op1(e11, e11)) = op2(h26(e11), h26(e11))) & (h26(op1(e11, e10)) = op2(h26(e11), h26(e10))) & (h26(op1(e10, e15)) = op2(h26(e10), h26(e15))) & (h26(op1(e10, e14)) = op2(h26(e10), h26(e14))) & (h26(op1(e10, e13)) = op2(h26(e10), h26(e13))) & (h26(op1(e10, e12)) = op2(h26(e10), h26(e12))) & (h26(op1(e10, e11)) = op2(h26(e10), h26(e11))) & (h26(op1(e10, e10)) = op2(h26(e10), h26(e10)))) | (((e25 = h25(e15)) | (e25 = h25(e14)) | (e25 = h25(e13)) | (e25 = h25(e12)) | (e25 = h25(e11)) | (e25 = h25(e10))) & ((e24 = h25(e15)) | (e24 = h25(e14)) | (e24 = h25(e13)) | (e24 = h25(e12)) | (e24 = h25(e11)) | (e24 = h25(e10))) & ((e23 = h25(e15)) | (e23 = h25(e14)) | (e23 = h25(e13)) | (e23 = h25(e12)) | (e23 = h25(e11)) | (e23 = h25(e10))) & ((e22 = h25(e15)) | (e22 = h25(e14)) | (e22 = h25(e13)) | (e22 = h25(e12)) | (e22 = h25(e11)) | (e22 = h25(e10))) & ((e21 = h25(e15)) | (e21 = h25(e14)) | (e21 = h25(e13)) | (e21 = h25(e12)) | (e21 = h25(e11)) | (e21 = h25(e10))) & ((e20 = h25(e15)) | (e20 = h25(e14)) | (e20 = h25(e13)) | (e20 = h25(e12)) | (e20 = h25(e11)) | (e20 = h25(e10))) & (h25(op1(e15, e15)) = op2(h25(e15), h25(e15))) & (h25(op1(e15, e14)) = op2(h25(e15), h25(e14))) & (h25(op1(e15, e13)) = op2(h25(e15), h25(e13))) & (h25(op1(e15, e12)) = op2(h25(e15), h25(e12))) & (h25(op1(e15, e11)) = op2(h25(e15), h25(e11))) & (h25(op1(e15, e10)) = op2(h25(e15), h25(e10))) & (h25(op1(e14, e15)) = op2(h25(e14), h25(e15))) & (h25(op1(e14, e14)) = op2(h25(e14), h25(e14))) & (h25(op1(e14, e13)) = op2(h25(e14), h25(e13))) & (h25(op1(e14, e12)) = op2(h25(e14), h25(e12))) & (h25(op1(e14, e11)) = op2(h25(e14), h25(e11))) & (h25(op1(e14, e10)) = op2(h25(e14), h25(e10))) & (h25(op1(e13, e15)) = op2(h25(e13), h25(e15))) & (h25(op1(e13, e14)) = op2(h25(e13), h25(e14))) & (h25(op1(e13, e13)) = op2(h25(e13), h25(e13))) & (h25(op1(e13, e12)) = op2(h25(e13), h25(e12))) & (h25(op1(e13, e11)) = op2(h25(e13), h25(e11))) & (h25(op1(e13, e10)) = op2(h25(e13), h25(e10))) & (h25(op1(e12, e15)) = op2(h25(e12), h25(e15))) & (h25(op1(e12, e14)) = op2(h25(e12), h25(e14))) & (h25(op1(e12, e13)) = op2(h25(e12), h25(e13))) & (h25(op1(e12, e12)) = op2(h25(e12), h25(e12))) & (h25(op1(e12, e11)) = op2(h25(e12), h25(e11))) & (h25(op1(e12, e10)) = op2(h25(e12), h25(e10))) & (h25(op1(e11, e15)) = op2(h25(e11), h25(e15))) & (h25(op1(e11, e14)) = op2(h25(e11), h25(e14))) & (h25(op1(e11, e13)) = op2(h25(e11), h25(e13))) & (h25(op1(e11, e12)) = op2(h25(e11), h25(e12))) & (h25(op1(e11, e11)) = op2(h25(e11), h25(e11))) & (h25(op1(e11, e10)) = op2(h25(e11), h25(e10))) & (h25(op1(e10, e15)) = op2(h25(e10), h25(e15))) & (h25(op1(e10, e14)) = op2(h25(e10), h25(e14))) & (h25(op1(e10, e13)) = op2(h25(e10), h25(e13))) & (h25(op1(e10, e12)) = op2(h25(e10), h25(e12))) & (h25(op1(e10, e11)) = op2(h25(e10), h25(e11))) & (h25(op1(e10, e10)) = op2(h25(e10), h25(e10)))) | (((e25 = h24(e15)) | (e25 = h24(e14)) | (e25 = h24(e13)) | (e25 = h24(e12)) | (e25 = h24(e11)) | (e25 = h24(e10))) & ((e24 = h24(e15)) | (e24 = h24(e14)) | (e24 = h24(e13)) | (e24 = h24(e12)) | (e24 = h24(e11)) | (e24 = h24(e10))) & ((e23 = h24(e15)) | (e23 = h24(e14)) | (e23 = h24(e13)) | (e23 = h24(e12)) | (e23 = h24(e11)) | (e23 = h24(e10))) & ((e22 = h24(e15)) | (e22 = h24(e14)) | (e22 = h24(e13)) | (e22 = h24(e12)) | (e22 = h24(e11)) | (e22 = h24(e10))) & ((e21 = h24(e15)) | (e21 = h24(e14)) | (e21 = h24(e13)) | (e21 = h24(e12)) | (e21 = h24(e11)) | (e21 = h24(e10))) & ((e20 = h24(e15)) | (e20 = h24(e14)) | (e20 = h24(e13)) | (e20 = h24(e12)) | (e20 = h24(e11)) | (e20 = h24(e10))) & (h24(op1(e15, e15)) = op2(h24(e15), h24(e15))) & (h24(op1(e15, e14)) = op2(h24(e15), h24(e14))) & (h24(op1(e15, e13)) = op2(h24(e15), h24(e13))) & (h24(op1(e15, e12)) = op2(h24(e15), h24(e12))) & (h24(op1(e15, e11)) = op2(h24(e15), h24(e11))) & (h24(op1(e15, e10)) = op2(h24(e15), h24(e10))) & (h24(op1(e14, e15)) = op2(h24(e14), h24(e15))) & (h24(op1(e14, e14)) = op2(h24(e14), h24(e14))) & (h24(op1(e14, e13)) = op2(h24(e14), h24(e13))) & (h24(op1(e14, e12)) = op2(h24(e14), h24(e12))) & (h24(op1(e14, e11)) = op2(h24(e14), h24(e11))) & (h24(op1(e14, e10)) = op2(h24(e14), h24(e10))) & (h24(op1(e13, e15)) = op2(h24(e13), h24(e15))) & (h24(op1(e13, e14)) = op2(h24(e13), h24(e14))) & (h24(op1(e13, e13)) = op2(h24(e13), h24(e13))) & (h24(op1(e13, e12)) = op2(h24(e13), h24(e12))) & (h24(op1(e13, e11)) = op2(h24(e13), h24(e11))) & (h24(op1(e13, e10)) = op2(h24(e13), h24(e10))) & (h24(op1(e12, e15)) = op2(h24(e12), h24(e15))) & (h24(op1(e12, e14)) = op2(h24(e12), h24(e14))) & (h24(op1(e12, e13)) = op2(h24(e12), h24(e13))) & (h24(op1(e12, e12)) = op2(h24(e12), h24(e12))) & (h24(op1(e12, e11)) = op2(h24(e12), h24(e11))) & (h24(op1(e12, e10)) = op2(h24(e12), h24(e10))) & (h24(op1(e11, e15)) = op2(h24(e11), h24(e15))) & (h24(op1(e11, e14)) = op2(h24(e11), h24(e14))) & (h24(op1(e11, e13)) = op2(h24(e11), h24(e13))) & (h24(op1(e11, e12)) = op2(h24(e11), h24(e12))) & (h24(op1(e11, e11)) = op2(h24(e11), h24(e11))) & (h24(op1(e11, e10)) = op2(h24(e11), h24(e10))) & (h24(op1(e10, e15)) = op2(h24(e10), h24(e15))) & (h24(op1(e10, e14)) = op2(h24(e10), h24(e14))) & (h24(op1(e10, e13)) = op2(h24(e10), h24(e13))) & (h24(op1(e10, e12)) = op2(h24(e10), h24(e12))) & (h24(op1(e10, e11)) = op2(h24(e10), h24(e11))) & (h24(op1(e10, e10)) = op2(h24(e10), h24(e10)))) | (((e25 = h23(e15)) | (e25 = h23(e14)) | (e25 = h23(e13)) | (e25 = h23(e12)) | (e25 = h23(e11)) | (e25 = h23(e10))) & ((e24 = h23(e15)) | (e24 = h23(e14)) | (e24 = h23(e13)) | (e24 = h23(e12)) | (e24 = h23(e11)) | (e24 = h23(e10))) & ((e23 = h23(e15)) | (e23 = h23(e14)) | (e23 = h23(e13)) | (e23 = h23(e12)) | (e23 = h23(e11)) | (e23 = h23(e10))) & ((e22 = h23(e15)) | (e22 = h23(e14)) | (e22 = h23(e13)) | (e22 = h23(e12)) | (e22 = h23(e11)) | (e22 = h23(e10))) & ((e21 = h23(e15)) | (e21 = h23(e14)) | (e21 = h23(e13)) | (e21 = h23(e12)) | (e21 = h23(e11)) | (e21 = h23(e10))) & ((e20 = h23(e15)) | (e20 = h23(e14)) | (e20 = h23(e13)) | (e20 = h23(e12)) | (e20 = h23(e11)) | (e20 = h23(e10))) & (h23(op1(e15, e15)) = op2(h23(e15), h23(e15))) & (h23(op1(e15, e14)) = op2(h23(e15), h23(e14))) & (h23(op1(e15, e13)) = op2(h23(e15), h23(e13))) & (h23(op1(e15, e12)) = op2(h23(e15), h23(e12))) & (h23(op1(e15, e11)) = op2(h23(e15), h23(e11))) & (h23(op1(e15, e10)) = op2(h23(e15), h23(e10))) & (h23(op1(e14, e15)) = op2(h23(e14), h23(e15))) & (h23(op1(e14, e14)) = op2(h23(e14), h23(e14))) & (h23(op1(e14, e13)) = op2(h23(e14), h23(e13))) & (h23(op1(e14, e12)) = op2(h23(e14), h23(e12))) & (h23(op1(e14, e11)) = op2(h23(e14), h23(e11))) & (h23(op1(e14, e10)) = op2(h23(e14), h23(e10))) & (h23(op1(e13, e15)) = op2(h23(e13), h23(e15))) & (h23(op1(e13, e14)) = op2(h23(e13), h23(e14))) & (h23(op1(e13, e13)) = op2(h23(e13), h23(e13))) & (h23(op1(e13, e12)) = op2(h23(e13), h23(e12))) & (h23(op1(e13, e11)) = op2(h23(e13), h23(e11))) & (h23(op1(e13, e10)) = op2(h23(e13), h23(e10))) & (h23(op1(e12, e15)) = op2(h23(e12), h23(e15))) & (h23(op1(e12, e14)) = op2(h23(e12), h23(e14))) & (h23(op1(e12, e13)) = op2(h23(e12), h23(e13))) & (h23(op1(e12, e12)) = op2(h23(e12), h23(e12))) & (h23(op1(e12, e11)) = op2(h23(e12), h23(e11))) & (h23(op1(e12, e10)) = op2(h23(e12), h23(e10))) & (h23(op1(e11, e15)) = op2(h23(e11), h23(e15))) & (h23(op1(e11, e14)) = op2(h23(e11), h23(e14))) & (h23(op1(e11, e13)) = op2(h23(e11), h23(e13))) & (h23(op1(e11, e12)) = op2(h23(e11), h23(e12))) & (h23(op1(e11, e11)) = op2(h23(e11), h23(e11))) & (h23(op1(e11, e10)) = op2(h23(e11), h23(e10))) & (h23(op1(e10, e15)) = op2(h23(e10), h23(e15))) & (h23(op1(e10, e14)) = op2(h23(e10), h23(e14))) & (h23(op1(e10, e13)) = op2(h23(e10), h23(e13))) & (h23(op1(e10, e12)) = op2(h23(e10), h23(e12))) & (h23(op1(e10, e11)) = op2(h23(e10), h23(e11))) & (h23(op1(e10, e10)) = op2(h23(e10), h23(e10)))) | (((e25 = h22(e15)) | (e25 = h22(e14)) | (e25 = h22(e13)) | (e25 = h22(e12)) | (e25 = h22(e11)) | (e25 = h22(e10))) & ((e24 = h22(e15)) | (e24 = h22(e14)) | (e24 = h22(e13)) | (e24 = h22(e12)) | (e24 = h22(e11)) | (e24 = h22(e10))) & ((e23 = h22(e15)) | (e23 = h22(e14)) | (e23 = h22(e13)) | (e23 = h22(e12)) | (e23 = h22(e11)) | (e23 = h22(e10))) & ((e22 = h22(e15)) | (e22 = h22(e14)) | (e22 = h22(e13)) | (e22 = h22(e12)) | (e22 = h22(e11)) | (e22 = h22(e10))) & ((e21 = h22(e15)) | (e21 = h22(e14)) | (e21 = h22(e13)) | (e21 = h22(e12)) | (e21 = h22(e11)) | (e21 = h22(e10))) & ((e20 = h22(e15)) | (e20 = h22(e14)) | (e20 = h22(e13)) | (e20 = h22(e12)) | (e20 = h22(e11)) | (e20 = h22(e10))) & (h22(op1(e15, e15)) = op2(h22(e15), h22(e15))) & (h22(op1(e15, e14)) = op2(h22(e15), h22(e14))) & (h22(op1(e15, e13)) = op2(h22(e15), h22(e13))) & (h22(op1(e15, e12)) = op2(h22(e15), h22(e12))) & (h22(op1(e15, e11)) = op2(h22(e15), h22(e11))) & (h22(op1(e15, e10)) = op2(h22(e15), h22(e10))) & (h22(op1(e14, e15)) = op2(h22(e14), h22(e15))) & (h22(op1(e14, e14)) = op2(h22(e14), h22(e14))) & (h22(op1(e14, e13)) = op2(h22(e14), h22(e13))) & (h22(op1(e14, e12)) = op2(h22(e14), h22(e12))) & (h22(op1(e14, e11)) = op2(h22(e14), h22(e11))) & (h22(op1(e14, e10)) = op2(h22(e14), h22(e10))) & (h22(op1(e13, e15)) = op2(h22(e13), h22(e15))) & (h22(op1(e13, e14)) = op2(h22(e13), h22(e14))) & (h22(op1(e13, e13)) = op2(h22(e13), h22(e13))) & (h22(op1(e13, e12)) = op2(h22(e13), h22(e12))) & (h22(op1(e13, e11)) = op2(h22(e13), h22(e11))) & (h22(op1(e13, e10)) = op2(h22(e13), h22(e10))) & (h22(op1(e12, e15)) = op2(h22(e12), h22(e15))) & (h22(op1(e12, e14)) = op2(h22(e12), h22(e14))) & (h22(op1(e12, e13)) = op2(h22(e12), h22(e13))) & (h22(op1(e12, e12)) = op2(h22(e12), h22(e12))) & (h22(op1(e12, e11)) = op2(h22(e12), h22(e11))) & (h22(op1(e12, e10)) = op2(h22(e12), h22(e10))) & (h22(op1(e11, e15)) = op2(h22(e11), h22(e15))) & (h22(op1(e11, e14)) = op2(h22(e11), h22(e14))) & (h22(op1(e11, e13)) = op2(h22(e11), h22(e13))) & (h22(op1(e11, e12)) = op2(h22(e11), h22(e12))) & (h22(op1(e11, e11)) = op2(h22(e11), h22(e11))) & (h22(op1(e11, e10)) = op2(h22(e11), h22(e10))) & (h22(op1(e10, e15)) = op2(h22(e10), h22(e15))) & (h22(op1(e10, e14)) = op2(h22(e10), h22(e14))) & (h22(op1(e10, e13)) = op2(h22(e10), h22(e13))) & (h22(op1(e10, e12)) = op2(h22(e10), h22(e12))) & (h22(op1(e10, e11)) = op2(h22(e10), h22(e11))) & (h22(op1(e10, e10)) = op2(h22(e10), h22(e10)))) | (((e25 = h21(e15)) | (e25 = h21(e14)) | (e25 = h21(e13)) | (e25 = h21(e12)) | (e25 = h21(e11)) | (e25 = h21(e10))) & ((e24 = h21(e15)) | (e24 = h21(e14)) | (e24 = h21(e13)) | (e24 = h21(e12)) | (e24 = h21(e11)) | (e24 = h21(e10))) & ((e23 = h21(e15)) | (e23 = h21(e14)) | (e23 = h21(e13)) | (e23 = h21(e12)) | (e23 = h21(e11)) | (e23 = h21(e10))) & ((e22 = h21(e15)) | (e22 = h21(e14)) | (e22 = h21(e13)) | (e22 = h21(e12)) | (e22 = h21(e11)) | (e22 = h21(e10))) & ((e21 = h21(e15)) | (e21 = h21(e14)) | (e21 = h21(e13)) | (e21 = h21(e12)) | (e21 = h21(e11)) | (e21 = h21(e10))) & ((e20 = h21(e15)) | (e20 = h21(e14)) | (e20 = h21(e13)) | (e20 = h21(e12)) | (e20 = h21(e11)) | (e20 = h21(e10))) & (h21(op1(e15, e15)) = op2(h21(e15), h21(e15))) & (h21(op1(e15, e14)) = op2(h21(e15), h21(e14))) & (h21(op1(e15, e13)) = op2(h21(e15), h21(e13))) & (h21(op1(e15, e12)) = op2(h21(e15), h21(e12))) & (h21(op1(e15, e11)) = op2(h21(e15), h21(e11))) & (h21(op1(e15, e10)) = op2(h21(e15), h21(e10))) & (h21(op1(e14, e15)) = op2(h21(e14), h21(e15))) & (h21(op1(e14, e14)) = op2(h21(e14), h21(e14))) & (h21(op1(e14, e13)) = op2(h21(e14), h21(e13))) & (h21(op1(e14, e12)) = op2(h21(e14), h21(e12))) & (h21(op1(e14, e11)) = op2(h21(e14), h21(e11))) & (h21(op1(e14, e10)) = op2(h21(e14), h21(e10))) & (h21(op1(e13, e15)) = op2(h21(e13), h21(e15))) & (h21(op1(e13, e14)) = op2(h21(e13), h21(e14))) & (h21(op1(e13, e13)) = op2(h21(e13), h21(e13))) & (h21(op1(e13, e12)) = op2(h21(e13), h21(e12))) & (h21(op1(e13, e11)) = op2(h21(e13), h21(e11))) & (h21(op1(e13, e10)) = op2(h21(e13), h21(e10))) & (h21(op1(e12, e15)) = op2(h21(e12), h21(e15))) & (h21(op1(e12, e14)) = op2(h21(e12), h21(e14))) & (h21(op1(e12, e13)) = op2(h21(e12), h21(e13))) & (h21(op1(e12, e12)) = op2(h21(e12), h21(e12))) & (h21(op1(e12, e11)) = op2(h21(e12), h21(e11))) & (h21(op1(e12, e10)) = op2(h21(e12), h21(e10))) & (h21(op1(e11, e15)) = op2(h21(e11), h21(e15))) & (h21(op1(e11, e14)) = op2(h21(e11), h21(e14))) & (h21(op1(e11, e13)) = op2(h21(e11), h21(e13))) & (h21(op1(e11, e12)) = op2(h21(e11), h21(e12))) & (h21(op1(e11, e11)) = op2(h21(e11), h21(e11))) & (h21(op1(e11, e10)) = op2(h21(e11), h21(e10))) & (h21(op1(e10, e15)) = op2(h21(e10), h21(e15))) & (h21(op1(e10, e14)) = op2(h21(e10), h21(e14))) & (h21(op1(e10, e13)) = op2(h21(e10), h21(e13))) & (h21(op1(e10, e12)) = op2(h21(e10), h21(e12))) & (h21(op1(e10, e11)) = op2(h21(e10), h21(e11))) & (h21(op1(e10, e10)) = op2(h21(e10), h21(e10)))) | (((e25 = h20(e15)) | (e25 = h20(e14)) | (e25 = h20(e13)) | (e25 = h20(e12)) | (e25 = h20(e11)) | (e25 = h20(e10))) & ((e24 = h20(e15)) | (e24 = h20(e14)) | (e24 = h20(e13)) | (e24 = h20(e12)) | (e24 = h20(e11)) | (e24 = h20(e10))) & ((e23 = h20(e15)) | (e23 = h20(e14)) | (e23 = h20(e13)) | (e23 = h20(e12)) | (e23 = h20(e11)) | (e23 = h20(e10))) & ((e22 = h20(e15)) | (e22 = h20(e14)) | (e22 = h20(e13)) | (e22 = h20(e12)) | (e22 = h20(e11)) | (e22 = h20(e10))) & ((e21 = h20(e15)) | (e21 = h20(e14)) | (e21 = h20(e13)) | (e21 = h20(e12)) | (e21 = h20(e11)) | (e21 = h20(e10))) & ((e20 = h20(e15)) | (e20 = h20(e14)) | (e20 = h20(e13)) | (e20 = h20(e12)) | (e20 = h20(e11)) | (e20 = h20(e10))) & (h20(op1(e15, e15)) = op2(h20(e15), h20(e15))) & (h20(op1(e15, e14)) = op2(h20(e15), h20(e14))) & (h20(op1(e15, e13)) = op2(h20(e15), h20(e13))) & (h20(op1(e15, e12)) = op2(h20(e15), h20(e12))) & (h20(op1(e15, e11)) = op2(h20(e15), h20(e11))) & (h20(op1(e15, e10)) = op2(h20(e15), h20(e10))) & (h20(op1(e14, e15)) = op2(h20(e14), h20(e15))) & (h20(op1(e14, e14)) = op2(h20(e14), h20(e14))) & (h20(op1(e14, e13)) = op2(h20(e14), h20(e13))) & (h20(op1(e14, e12)) = op2(h20(e14), h20(e12))) & (h20(op1(e14, e11)) = op2(h20(e14), h20(e11))) & (h20(op1(e14, e10)) = op2(h20(e14), h20(e10))) & (h20(op1(e13, e15)) = op2(h20(e13), h20(e15))) & (h20(op1(e13, e14)) = op2(h20(e13), h20(e14))) & (h20(op1(e13, e13)) = op2(h20(e13), h20(e13))) & (h20(op1(e13, e12)) = op2(h20(e13), h20(e12))) & (h20(op1(e13, e11)) = op2(h20(e13), h20(e11))) & (h20(op1(e13, e10)) = op2(h20(e13), h20(e10))) & (h20(op1(e12, e15)) = op2(h20(e12), h20(e15))) & (h20(op1(e12, e14)) = op2(h20(e12), h20(e14))) & (h20(op1(e12, e13)) = op2(h20(e12), h20(e13))) & (h20(op1(e12, e12)) = op2(h20(e12), h20(e12))) & (h20(op1(e12, e11)) = op2(h20(e12), h20(e11))) & (h20(op1(e12, e10)) = op2(h20(e12), h20(e10))) & (h20(op1(e11, e15)) = op2(h20(e11), h20(e15))) & (h20(op1(e11, e14)) = op2(h20(e11), h20(e14))) & (h20(op1(e11, e13)) = op2(h20(e11), h20(e13))) & (h20(op1(e11, e12)) = op2(h20(e11), h20(e12))) & (h20(op1(e11, e11)) = op2(h20(e11), h20(e11))) & (h20(op1(e11, e10)) = op2(h20(e11), h20(e10))) & (h20(op1(e10, e15)) = op2(h20(e10), h20(e15))) & (h20(op1(e10, e14)) = op2(h20(e10), h20(e14))) & (h20(op1(e10, e13)) = op2(h20(e10), h20(e13))) & (h20(op1(e10, e12)) = op2(h20(e10), h20(e12))) & (h20(op1(e10, e11)) = op2(h20(e10), h20(e11))) & (h20(op1(e10, e10)) = op2(h20(e10), h20(e10)))) | (((e25 = h19(e15)) | (e25 = h19(e14)) | (e25 = h19(e13)) | (e25 = h19(e12)) | (e25 = h19(e11)) | (e25 = h19(e10))) & ((e24 = h19(e15)) | (e24 = h19(e14)) | (e24 = h19(e13)) | (e24 = h19(e12)) | (e24 = h19(e11)) | (e24 = h19(e10))) & ((e23 = h19(e15)) | (e23 = h19(e14)) | (e23 = h19(e13)) | (e23 = h19(e12)) | (e23 = h19(e11)) | (e23 = h19(e10))) & ((e22 = h19(e15)) | (e22 = h19(e14)) | (e22 = h19(e13)) | (e22 = h19(e12)) | (e22 = h19(e11)) | (e22 = h19(e10))) & ((e21 = h19(e15)) | (e21 = h19(e14)) | (e21 = h19(e13)) | (e21 = h19(e12)) | (e21 = h19(e11)) | (e21 = h19(e10))) & ((e20 = h19(e15)) | (e20 = h19(e14)) | (e20 = h19(e13)) | (e20 = h19(e12)) | (e20 = h19(e11)) | (e20 = h19(e10))) & (h19(op1(e15, e15)) = op2(h19(e15), h19(e15))) & (h19(op1(e15, e14)) = op2(h19(e15), h19(e14))) & (h19(op1(e15, e13)) = op2(h19(e15), h19(e13))) & (h19(op1(e15, e12)) = op2(h19(e15), h19(e12))) & (h19(op1(e15, e11)) = op2(h19(e15), h19(e11))) & (h19(op1(e15, e10)) = op2(h19(e15), h19(e10))) & (h19(op1(e14, e15)) = op2(h19(e14), h19(e15))) & (h19(op1(e14, e14)) = op2(h19(e14), h19(e14))) & (h19(op1(e14, e13)) = op2(h19(e14), h19(e13))) & (h19(op1(e14, e12)) = op2(h19(e14), h19(e12))) & (h19(op1(e14, e11)) = op2(h19(e14), h19(e11))) & (h19(op1(e14, e10)) = op2(h19(e14), h19(e10))) & (h19(op1(e13, e15)) = op2(h19(e13), h19(e15))) & (h19(op1(e13, e14)) = op2(h19(e13), h19(e14))) & (h19(op1(e13, e13)) = op2(h19(e13), h19(e13))) & (h19(op1(e13, e12)) = op2(h19(e13), h19(e12))) & (h19(op1(e13, e11)) = op2(h19(e13), h19(e11))) & (h19(op1(e13, e10)) = op2(h19(e13), h19(e10))) & (h19(op1(e12, e15)) = op2(h19(e12), h19(e15))) & (h19(op1(e12, e14)) = op2(h19(e12), h19(e14))) & (h19(op1(e12, e13)) = op2(h19(e12), h19(e13))) & (h19(op1(e12, e12)) = op2(h19(e12), h19(e12))) & (h19(op1(e12, e11)) = op2(h19(e12), h19(e11))) & (h19(op1(e12, e10)) = op2(h19(e12), h19(e10))) & (h19(op1(e11, e15)) = op2(h19(e11), h19(e15))) & (h19(op1(e11, e14)) = op2(h19(e11), h19(e14))) & (h19(op1(e11, e13)) = op2(h19(e11), h19(e13))) & (h19(op1(e11, e12)) = op2(h19(e11), h19(e12))) & (h19(op1(e11, e11)) = op2(h19(e11), h19(e11))) & (h19(op1(e11, e10)) = op2(h19(e11), h19(e10))) & (h19(op1(e10, e15)) = op2(h19(e10), h19(e15))) & (h19(op1(e10, e14)) = op2(h19(e10), h19(e14))) & (h19(op1(e10, e13)) = op2(h19(e10), h19(e13))) & (h19(op1(e10, e12)) = op2(h19(e10), h19(e12))) & (h19(op1(e10, e11)) = op2(h19(e10), h19(e11))) & (h19(op1(e10, e10)) = op2(h19(e10), h19(e10)))) | (((e25 = h18(e15)) | (e25 = h18(e14)) | (e25 = h18(e13)) | (e25 = h18(e12)) | (e25 = h18(e11)) | (e25 = h18(e10))) & ((e24 = h18(e15)) | (e24 = h18(e14)) | (e24 = h18(e13)) | (e24 = h18(e12)) | (e24 = h18(e11)) | (e24 = h18(e10))) & ((e23 = h18(e15)) | (e23 = h18(e14)) | (e23 = h18(e13)) | (e23 = h18(e12)) | (e23 = h18(e11)) | (e23 = h18(e10))) & ((e22 = h18(e15)) | (e22 = h18(e14)) | (e22 = h18(e13)) | (e22 = h18(e12)) | (e22 = h18(e11)) | (e22 = h18(e10))) & ((e21 = h18(e15)) | (e21 = h18(e14)) | (e21 = h18(e13)) | (e21 = h18(e12)) | (e21 = h18(e11)) | (e21 = h18(e10))) & ((e20 = h18(e15)) | (e20 = h18(e14)) | (e20 = h18(e13)) | (e20 = h18(e12)) | (e20 = h18(e11)) | (e20 = h18(e10))) & (h18(op1(e15, e15)) = op2(h18(e15), h18(e15))) & (h18(op1(e15, e14)) = op2(h18(e15), h18(e14))) & (h18(op1(e15, e13)) = op2(h18(e15), h18(e13))) & (h18(op1(e15, e12)) = op2(h18(e15), h18(e12))) & (h18(op1(e15, e11)) = op2(h18(e15), h18(e11))) & (h18(op1(e15, e10)) = op2(h18(e15), h18(e10))) & (h18(op1(e14, e15)) = op2(h18(e14), h18(e15))) & (h18(op1(e14, e14)) = op2(h18(e14), h18(e14))) & (h18(op1(e14, e13)) = op2(h18(e14), h18(e13))) & (h18(op1(e14, e12)) = op2(h18(e14), h18(e12))) & (h18(op1(e14, e11)) = op2(h18(e14), h18(e11))) & (h18(op1(e14, e10)) = op2(h18(e14), h18(e10))) & (h18(op1(e13, e15)) = op2(h18(e13), h18(e15))) & (h18(op1(e13, e14)) = op2(h18(e13), h18(e14))) & (h18(op1(e13, e13)) = op2(h18(e13), h18(e13))) & (h18(op1(e13, e12)) = op2(h18(e13), h18(e12))) & (h18(op1(e13, e11)) = op2(h18(e13), h18(e11))) & (h18(op1(e13, e10)) = op2(h18(e13), h18(e10))) & (h18(op1(e12, e15)) = op2(h18(e12), h18(e15))) & (h18(op1(e12, e14)) = op2(h18(e12), h18(e14))) & (h18(op1(e12, e13)) = op2(h18(e12), h18(e13))) & (h18(op1(e12, e12)) = op2(h18(e12), h18(e12))) & (h18(op1(e12, e11)) = op2(h18(e12), h18(e11))) & (h18(op1(e12, e10)) = op2(h18(e12), h18(e10))) & (h18(op1(e11, e15)) = op2(h18(e11), h18(e15))) & (h18(op1(e11, e14)) = op2(h18(e11), h18(e14))) & (h18(op1(e11, e13)) = op2(h18(e11), h18(e13))) & (h18(op1(e11, e12)) = op2(h18(e11), h18(e12))) & (h18(op1(e11, e11)) = op2(h18(e11), h18(e11))) & (h18(op1(e11, e10)) = op2(h18(e11), h18(e10))) & (h18(op1(e10, e15)) = op2(h18(e10), h18(e15))) & (h18(op1(e10, e14)) = op2(h18(e10), h18(e14))) & (h18(op1(e10, e13)) = op2(h18(e10), h18(e13))) & (h18(op1(e10, e12)) = op2(h18(e10), h18(e12))) & (h18(op1(e10, e11)) = op2(h18(e10), h18(e11))) & (h18(op1(e10, e10)) = op2(h18(e10), h18(e10)))) | (((e25 = h17(e15)) | (e25 = h17(e14)) | (e25 = h17(e13)) | (e25 = h17(e12)) | (e25 = h17(e11)) | (e25 = h17(e10))) & ((e24 = h17(e15)) | (e24 = h17(e14)) | (e24 = h17(e13)) | (e24 = h17(e12)) | (e24 = h17(e11)) | (e24 = h17(e10))) & ((e23 = h17(e15)) | (e23 = h17(e14)) | (e23 = h17(e13)) | (e23 = h17(e12)) | (e23 = h17(e11)) | (e23 = h17(e10))) & ((e22 = h17(e15)) | (e22 = h17(e14)) | (e22 = h17(e13)) | (e22 = h17(e12)) | (e22 = h17(e11)) | (e22 = h17(e10))) & ((e21 = h17(e15)) | (e21 = h17(e14)) | (e21 = h17(e13)) | (e21 = h17(e12)) | (e21 = h17(e11)) | (e21 = h17(e10))) & ((e20 = h17(e15)) | (e20 = h17(e14)) | (e20 = h17(e13)) | (e20 = h17(e12)) | (e20 = h17(e11)) | (e20 = h17(e10))) & (h17(op1(e15, e15)) = op2(h17(e15), h17(e15))) & (h17(op1(e15, e14)) = op2(h17(e15), h17(e14))) & (h17(op1(e15, e13)) = op2(h17(e15), h17(e13))) & (h17(op1(e15, e12)) = op2(h17(e15), h17(e12))) & (h17(op1(e15, e11)) = op2(h17(e15), h17(e11))) & (h17(op1(e15, e10)) = op2(h17(e15), h17(e10))) & (h17(op1(e14, e15)) = op2(h17(e14), h17(e15))) & (h17(op1(e14, e14)) = op2(h17(e14), h17(e14))) & (h17(op1(e14, e13)) = op2(h17(e14), h17(e13))) & (h17(op1(e14, e12)) = op2(h17(e14), h17(e12))) & (h17(op1(e14, e11)) = op2(h17(e14), h17(e11))) & (h17(op1(e14, e10)) = op2(h17(e14), h17(e10))) & (h17(op1(e13, e15)) = op2(h17(e13), h17(e15))) & (h17(op1(e13, e14)) = op2(h17(e13), h17(e14))) & (h17(op1(e13, e13)) = op2(h17(e13), h17(e13))) & (h17(op1(e13, e12)) = op2(h17(e13), h17(e12))) & (h17(op1(e13, e11)) = op2(h17(e13), h17(e11))) & (h17(op1(e13, e10)) = op2(h17(e13), h17(e10))) & (h17(op1(e12, e15)) = op2(h17(e12), h17(e15))) & (h17(op1(e12, e14)) = op2(h17(e12), h17(e14))) & (h17(op1(e12, e13)) = op2(h17(e12), h17(e13))) & (h17(op1(e12, e12)) = op2(h17(e12), h17(e12))) & (h17(op1(e12, e11)) = op2(h17(e12), h17(e11))) & (h17(op1(e12, e10)) = op2(h17(e12), h17(e10))) & (h17(op1(e11, e15)) = op2(h17(e11), h17(e15))) & (h17(op1(e11, e14)) = op2(h17(e11), h17(e14))) & (h17(op1(e11, e13)) = op2(h17(e11), h17(e13))) & (h17(op1(e11, e12)) = op2(h17(e11), h17(e12))) & (h17(op1(e11, e11)) = op2(h17(e11), h17(e11))) & (h17(op1(e11, e10)) = op2(h17(e11), h17(e10))) & (h17(op1(e10, e15)) = op2(h17(e10), h17(e15))) & (h17(op1(e10, e14)) = op2(h17(e10), h17(e14))) & (h17(op1(e10, e13)) = op2(h17(e10), h17(e13))) & (h17(op1(e10, e12)) = op2(h17(e10), h17(e12))) & (h17(op1(e10, e11)) = op2(h17(e10), h17(e11))) & (h17(op1(e10, e10)) = op2(h17(e10), h17(e10)))) | (((e25 = h16(e15)) | (e25 = h16(e14)) | (e25 = h16(e13)) | (e25 = h16(e12)) | (e25 = h16(e11)) | (e25 = h16(e10))) & ((e24 = h16(e15)) | (e24 = h16(e14)) | (e24 = h16(e13)) | (e24 = h16(e12)) | (e24 = h16(e11)) | (e24 = h16(e10))) & ((e23 = h16(e15)) | (e23 = h16(e14)) | (e23 = h16(e13)) | (e23 = h16(e12)) | (e23 = h16(e11)) | (e23 = h16(e10))) & ((e22 = h16(e15)) | (e22 = h16(e14)) | (e22 = h16(e13)) | (e22 = h16(e12)) | (e22 = h16(e11)) | (e22 = h16(e10))) & ((e21 = h16(e15)) | (e21 = h16(e14)) | (e21 = h16(e13)) | (e21 = h16(e12)) | (e21 = h16(e11)) | (e21 = h16(e10))) & ((e20 = h16(e15)) | (e20 = h16(e14)) | (e20 = h16(e13)) | (e20 = h16(e12)) | (e20 = h16(e11)) | (e20 = h16(e10))) & (h16(op1(e15, e15)) = op2(h16(e15), h16(e15))) & (h16(op1(e15, e14)) = op2(h16(e15), h16(e14))) & (h16(op1(e15, e13)) = op2(h16(e15), h16(e13))) & (h16(op1(e15, e12)) = op2(h16(e15), h16(e12))) & (h16(op1(e15, e11)) = op2(h16(e15), h16(e11))) & (h16(op1(e15, e10)) = op2(h16(e15), h16(e10))) & (h16(op1(e14, e15)) = op2(h16(e14), h16(e15))) & (h16(op1(e14, e14)) = op2(h16(e14), h16(e14))) & (h16(op1(e14, e13)) = op2(h16(e14), h16(e13))) & (h16(op1(e14, e12)) = op2(h16(e14), h16(e12))) & (h16(op1(e14, e11)) = op2(h16(e14), h16(e11))) & (h16(op1(e14, e10)) = op2(h16(e14), h16(e10))) & (h16(op1(e13, e15)) = op2(h16(e13), h16(e15))) & (h16(op1(e13, e14)) = op2(h16(e13), h16(e14))) & (h16(op1(e13, e13)) = op2(h16(e13), h16(e13))) & (h16(op1(e13, e12)) = op2(h16(e13), h16(e12))) & (h16(op1(e13, e11)) = op2(h16(e13), h16(e11))) & (h16(op1(e13, e10)) = op2(h16(e13), h16(e10))) & (h16(op1(e12, e15)) = op2(h16(e12), h16(e15))) & (h16(op1(e12, e14)) = op2(h16(e12), h16(e14))) & (h16(op1(e12, e13)) = op2(h16(e12), h16(e13))) & (h16(op1(e12, e12)) = op2(h16(e12), h16(e12))) & (h16(op1(e12, e11)) = op2(h16(e12), h16(e11))) & (h16(op1(e12, e10)) = op2(h16(e12), h16(e10))) & (h16(op1(e11, e15)) = op2(h16(e11), h16(e15))) & (h16(op1(e11, e14)) = op2(h16(e11), h16(e14))) & (h16(op1(e11, e13)) = op2(h16(e11), h16(e13))) & (h16(op1(e11, e12)) = op2(h16(e11), h16(e12))) & (h16(op1(e11, e11)) = op2(h16(e11), h16(e11))) & (h16(op1(e11, e10)) = op2(h16(e11), h16(e10))) & (h16(op1(e10, e15)) = op2(h16(e10), h16(e15))) & (h16(op1(e10, e14)) = op2(h16(e10), h16(e14))) & (h16(op1(e10, e13)) = op2(h16(e10), h16(e13))) & (h16(op1(e10, e12)) = op2(h16(e10), h16(e12))) & (h16(op1(e10, e11)) = op2(h16(e10), h16(e11))) & (h16(op1(e10, e10)) = op2(h16(e10), h16(e10)))) | (((e25 = h15(e15)) | (e25 = h15(e14)) | (e25 = h15(e13)) | (e25 = h15(e12)) | (e25 = h15(e11)) | (e25 = h15(e10))) & ((e24 = h15(e15)) | (e24 = h15(e14)) | (e24 = h15(e13)) | (e24 = h15(e12)) | (e24 = h15(e11)) | (e24 = h15(e10))) & ((e23 = h15(e15)) | (e23 = h15(e14)) | (e23 = h15(e13)) | (e23 = h15(e12)) | (e23 = h15(e11)) | (e23 = h15(e10))) & ((e22 = h15(e15)) | (e22 = h15(e14)) | (e22 = h15(e13)) | (e22 = h15(e12)) | (e22 = h15(e11)) | (e22 = h15(e10))) & ((e21 = h15(e15)) | (e21 = h15(e14)) | (e21 = h15(e13)) | (e21 = h15(e12)) | (e21 = h15(e11)) | (e21 = h15(e10))) & ((e20 = h15(e15)) | (e20 = h15(e14)) | (e20 = h15(e13)) | (e20 = h15(e12)) | (e20 = h15(e11)) | (e20 = h15(e10))) & (h15(op1(e15, e15)) = op2(h15(e15), h15(e15))) & (h15(op1(e15, e14)) = op2(h15(e15), h15(e14))) & (h15(op1(e15, e13)) = op2(h15(e15), h15(e13))) & (h15(op1(e15, e12)) = op2(h15(e15), h15(e12))) & (h15(op1(e15, e11)) = op2(h15(e15), h15(e11))) & (h15(op1(e15, e10)) = op2(h15(e15), h15(e10))) & (h15(op1(e14, e15)) = op2(h15(e14), h15(e15))) & (h15(op1(e14, e14)) = op2(h15(e14), h15(e14))) & (h15(op1(e14, e13)) = op2(h15(e14), h15(e13))) & (h15(op1(e14, e12)) = op2(h15(e14), h15(e12))) & (h15(op1(e14, e11)) = op2(h15(e14), h15(e11))) & (h15(op1(e14, e10)) = op2(h15(e14), h15(e10))) & (h15(op1(e13, e15)) = op2(h15(e13), h15(e15))) & (h15(op1(e13, e14)) = op2(h15(e13), h15(e14))) & (h15(op1(e13, e13)) = op2(h15(e13), h15(e13))) & (h15(op1(e13, e12)) = op2(h15(e13), h15(e12))) & (h15(op1(e13, e11)) = op2(h15(e13), h15(e11))) & (h15(op1(e13, e10)) = op2(h15(e13), h15(e10))) & (h15(op1(e12, e15)) = op2(h15(e12), h15(e15))) & (h15(op1(e12, e14)) = op2(h15(e12), h15(e14))) & (h15(op1(e12, e13)) = op2(h15(e12), h15(e13))) & (h15(op1(e12, e12)) = op2(h15(e12), h15(e12))) & (h15(op1(e12, e11)) = op2(h15(e12), h15(e11))) & (h15(op1(e12, e10)) = op2(h15(e12), h15(e10))) & (h15(op1(e11, e15)) = op2(h15(e11), h15(e15))) & (h15(op1(e11, e14)) = op2(h15(e11), h15(e14))) & (h15(op1(e11, e13)) = op2(h15(e11), h15(e13))) & (h15(op1(e11, e12)) = op2(h15(e11), h15(e12))) & (h15(op1(e11, e11)) = op2(h15(e11), h15(e11))) & (h15(op1(e11, e10)) = op2(h15(e11), h15(e10))) & (h15(op1(e10, e15)) = op2(h15(e10), h15(e15))) & (h15(op1(e10, e14)) = op2(h15(e10), h15(e14))) & (h15(op1(e10, e13)) = op2(h15(e10), h15(e13))) & (h15(op1(e10, e12)) = op2(h15(e10), h15(e12))) & (h15(op1(e10, e11)) = op2(h15(e10), h15(e11))) & (h15(op1(e10, e10)) = op2(h15(e10), h15(e10)))) | (((e25 = h14(e15)) | (e25 = h14(e14)) | (e25 = h14(e13)) | (e25 = h14(e12)) | (e25 = h14(e11)) | (e25 = h14(e10))) & ((e24 = h14(e15)) | (e24 = h14(e14)) | (e24 = h14(e13)) | (e24 = h14(e12)) | (e24 = h14(e11)) | (e24 = h14(e10))) & ((e23 = h14(e15)) | (e23 = h14(e14)) | (e23 = h14(e13)) | (e23 = h14(e12)) | (e23 = h14(e11)) | (e23 = h14(e10))) & ((e22 = h14(e15)) | (e22 = h14(e14)) | (e22 = h14(e13)) | (e22 = h14(e12)) | (e22 = h14(e11)) | (e22 = h14(e10))) & ((e21 = h14(e15)) | (e21 = h14(e14)) | (e21 = h14(e13)) | (e21 = h14(e12)) | (e21 = h14(e11)) | (e21 = h14(e10))) & ((e20 = h14(e15)) | (e20 = h14(e14)) | (e20 = h14(e13)) | (e20 = h14(e12)) | (e20 = h14(e11)) | (e20 = h14(e10))) & (h14(op1(e15, e15)) = op2(h14(e15), h14(e15))) & (h14(op1(e15, e14)) = op2(h14(e15), h14(e14))) & (h14(op1(e15, e13)) = op2(h14(e15), h14(e13))) & (h14(op1(e15, e12)) = op2(h14(e15), h14(e12))) & (h14(op1(e15, e11)) = op2(h14(e15), h14(e11))) & (h14(op1(e15, e10)) = op2(h14(e15), h14(e10))) & (h14(op1(e14, e15)) = op2(h14(e14), h14(e15))) & (h14(op1(e14, e14)) = op2(h14(e14), h14(e14))) & (h14(op1(e14, e13)) = op2(h14(e14), h14(e13))) & (h14(op1(e14, e12)) = op2(h14(e14), h14(e12))) & (h14(op1(e14, e11)) = op2(h14(e14), h14(e11))) & (h14(op1(e14, e10)) = op2(h14(e14), h14(e10))) & (h14(op1(e13, e15)) = op2(h14(e13), h14(e15))) & (h14(op1(e13, e14)) = op2(h14(e13), h14(e14))) & (h14(op1(e13, e13)) = op2(h14(e13), h14(e13))) & (h14(op1(e13, e12)) = op2(h14(e13), h14(e12))) & (h14(op1(e13, e11)) = op2(h14(e13), h14(e11))) & (h14(op1(e13, e10)) = op2(h14(e13), h14(e10))) & (h14(op1(e12, e15)) = op2(h14(e12), h14(e15))) & (h14(op1(e12, e14)) = op2(h14(e12), h14(e14))) & (h14(op1(e12, e13)) = op2(h14(e12), h14(e13))) & (h14(op1(e12, e12)) = op2(h14(e12), h14(e12))) & (h14(op1(e12, e11)) = op2(h14(e12), h14(e11))) & (h14(op1(e12, e10)) = op2(h14(e12), h14(e10))) & (h14(op1(e11, e15)) = op2(h14(e11), h14(e15))) & (h14(op1(e11, e14)) = op2(h14(e11), h14(e14))) & (h14(op1(e11, e13)) = op2(h14(e11), h14(e13))) & (h14(op1(e11, e12)) = op2(h14(e11), h14(e12))) & (h14(op1(e11, e11)) = op2(h14(e11), h14(e11))) & (h14(op1(e11, e10)) = op2(h14(e11), h14(e10))) & (h14(op1(e10, e15)) = op2(h14(e10), h14(e15))) & (h14(op1(e10, e14)) = op2(h14(e10), h14(e14))) & (h14(op1(e10, e13)) = op2(h14(e10), h14(e13))) & (h14(op1(e10, e12)) = op2(h14(e10), h14(e12))) & (h14(op1(e10, e11)) = op2(h14(e10), h14(e11))) & (h14(op1(e10, e10)) = op2(h14(e10), h14(e10)))) | (((e25 = h13(e15)) | (e25 = h13(e14)) | (e25 = h13(e13)) | (e25 = h13(e12)) | (e25 = h13(e11)) | (e25 = h13(e10))) & ((e24 = h13(e15)) | (e24 = h13(e14)) | (e24 = h13(e13)) | (e24 = h13(e12)) | (e24 = h13(e11)) | (e24 = h13(e10))) & ((e23 = h13(e15)) | (e23 = h13(e14)) | (e23 = h13(e13)) | (e23 = h13(e12)) | (e23 = h13(e11)) | (e23 = h13(e10))) & ((e22 = h13(e15)) | (e22 = h13(e14)) | (e22 = h13(e13)) | (e22 = h13(e12)) | (e22 = h13(e11)) | (e22 = h13(e10))) & ((e21 = h13(e15)) | (e21 = h13(e14)) | (e21 = h13(e13)) | (e21 = h13(e12)) | (e21 = h13(e11)) | (e21 = h13(e10))) & ((e20 = h13(e15)) | (e20 = h13(e14)) | (e20 = h13(e13)) | (e20 = h13(e12)) | (e20 = h13(e11)) | (e20 = h13(e10))) & (h13(op1(e15, e15)) = op2(h13(e15), h13(e15))) & (h13(op1(e15, e14)) = op2(h13(e15), h13(e14))) & (h13(op1(e15, e13)) = op2(h13(e15), h13(e13))) & (h13(op1(e15, e12)) = op2(h13(e15), h13(e12))) & (h13(op1(e15, e11)) = op2(h13(e15), h13(e11))) & (h13(op1(e15, e10)) = op2(h13(e15), h13(e10))) & (h13(op1(e14, e15)) = op2(h13(e14), h13(e15))) & (h13(op1(e14, e14)) = op2(h13(e14), h13(e14))) & (h13(op1(e14, e13)) = op2(h13(e14), h13(e13))) & (h13(op1(e14, e12)) = op2(h13(e14), h13(e12))) & (h13(op1(e14, e11)) = op2(h13(e14), h13(e11))) & (h13(op1(e14, e10)) = op2(h13(e14), h13(e10))) & (h13(op1(e13, e15)) = op2(h13(e13), h13(e15))) & (h13(op1(e13, e14)) = op2(h13(e13), h13(e14))) & (h13(op1(e13, e13)) = op2(h13(e13), h13(e13))) & (h13(op1(e13, e12)) = op2(h13(e13), h13(e12))) & (h13(op1(e13, e11)) = op2(h13(e13), h13(e11))) & (h13(op1(e13, e10)) = op2(h13(e13), h13(e10))) & (h13(op1(e12, e15)) = op2(h13(e12), h13(e15))) & (h13(op1(e12, e14)) = op2(h13(e12), h13(e14))) & (h13(op1(e12, e13)) = op2(h13(e12), h13(e13))) & (h13(op1(e12, e12)) = op2(h13(e12), h13(e12))) & (h13(op1(e12, e11)) = op2(h13(e12), h13(e11))) & (h13(op1(e12, e10)) = op2(h13(e12), h13(e10))) & (h13(op1(e11, e15)) = op2(h13(e11), h13(e15))) & (h13(op1(e11, e14)) = op2(h13(e11), h13(e14))) & (h13(op1(e11, e13)) = op2(h13(e11), h13(e13))) & (h13(op1(e11, e12)) = op2(h13(e11), h13(e12))) & (h13(op1(e11, e11)) = op2(h13(e11), h13(e11))) & (h13(op1(e11, e10)) = op2(h13(e11), h13(e10))) & (h13(op1(e10, e15)) = op2(h13(e10), h13(e15))) & (h13(op1(e10, e14)) = op2(h13(e10), h13(e14))) & (h13(op1(e10, e13)) = op2(h13(e10), h13(e13))) & (h13(op1(e10, e12)) = op2(h13(e10), h13(e12))) & (h13(op1(e10, e11)) = op2(h13(e10), h13(e11))) & (h13(op1(e10, e10)) = op2(h13(e10), h13(e10)))) | (((e25 = h12(e15)) | (e25 = h12(e14)) | (e25 = h12(e13)) | (e25 = h12(e12)) | (e25 = h12(e11)) | (e25 = h12(e10))) & ((e24 = h12(e15)) | (e24 = h12(e14)) | (e24 = h12(e13)) | (e24 = h12(e12)) | (e24 = h12(e11)) | (e24 = h12(e10))) & ((e23 = h12(e15)) | (e23 = h12(e14)) | (e23 = h12(e13)) | (e23 = h12(e12)) | (e23 = h12(e11)) | (e23 = h12(e10))) & ((e22 = h12(e15)) | (e22 = h12(e14)) | (e22 = h12(e13)) | (e22 = h12(e12)) | (e22 = h12(e11)) | (e22 = h12(e10))) & ((e21 = h12(e15)) | (e21 = h12(e14)) | (e21 = h12(e13)) | (e21 = h12(e12)) | (e21 = h12(e11)) | (e21 = h12(e10))) & ((e20 = h12(e15)) | (e20 = h12(e14)) | (e20 = h12(e13)) | (e20 = h12(e12)) | (e20 = h12(e11)) | (e20 = h12(e10))) & (h12(op1(e15, e15)) = op2(h12(e15), h12(e15))) & (h12(op1(e15, e14)) = op2(h12(e15), h12(e14))) & (h12(op1(e15, e13)) = op2(h12(e15), h12(e13))) & (h12(op1(e15, e12)) = op2(h12(e15), h12(e12))) & (h12(op1(e15, e11)) = op2(h12(e15), h12(e11))) & (h12(op1(e15, e10)) = op2(h12(e15), h12(e10))) & (h12(op1(e14, e15)) = op2(h12(e14), h12(e15))) & (h12(op1(e14, e14)) = op2(h12(e14), h12(e14))) & (h12(op1(e14, e13)) = op2(h12(e14), h12(e13))) & (h12(op1(e14, e12)) = op2(h12(e14), h12(e12))) & (h12(op1(e14, e11)) = op2(h12(e14), h12(e11))) & (h12(op1(e14, e10)) = op2(h12(e14), h12(e10))) & (h12(op1(e13, e15)) = op2(h12(e13), h12(e15))) & (h12(op1(e13, e14)) = op2(h12(e13), h12(e14))) & (h12(op1(e13, e13)) = op2(h12(e13), h12(e13))) & (h12(op1(e13, e12)) = op2(h12(e13), h12(e12))) & (h12(op1(e13, e11)) = op2(h12(e13), h12(e11))) & (h12(op1(e13, e10)) = op2(h12(e13), h12(e10))) & (h12(op1(e12, e15)) = op2(h12(e12), h12(e15))) & (h12(op1(e12, e14)) = op2(h12(e12), h12(e14))) & (h12(op1(e12, e13)) = op2(h12(e12), h12(e13))) & (h12(op1(e12, e12)) = op2(h12(e12), h12(e12))) & (h12(op1(e12, e11)) = op2(h12(e12), h12(e11))) & (h12(op1(e12, e10)) = op2(h12(e12), h12(e10))) & (h12(op1(e11, e15)) = op2(h12(e11), h12(e15))) & (h12(op1(e11, e14)) = op2(h12(e11), h12(e14))) & (h12(op1(e11, e13)) = op2(h12(e11), h12(e13))) & (h12(op1(e11, e12)) = op2(h12(e11), h12(e12))) & (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) & (h12(op1(e11, e10)) = op2(h12(e11), h12(e10))) & (h12(op1(e10, e15)) = op2(h12(e10), h12(e15))) & (h12(op1(e10, e14)) = op2(h12(e10), h12(e14))) & (h12(op1(e10, e13)) = op2(h12(e10), h12(e13))) & (h12(op1(e10, e12)) = op2(h12(e10), h12(e12))) & (h12(op1(e10, e11)) = op2(h12(e10), h12(e11))) & (h12(op1(e10, e10)) = op2(h12(e10), h12(e10)))) | (((e25 = h11(e15)) | (e25 = h11(e14)) | (e25 = h11(e13)) | (e25 = h11(e12)) | (e25 = h11(e11)) | (e25 = h11(e10))) & ((e24 = h11(e15)) | (e24 = h11(e14)) | (e24 = h11(e13)) | (e24 = h11(e12)) | (e24 = h11(e11)) | (e24 = h11(e10))) & ((e23 = h11(e15)) | (e23 = h11(e14)) | (e23 = h11(e13)) | (e23 = h11(e12)) | (e23 = h11(e11)) | (e23 = h11(e10))) & ((e22 = h11(e15)) | (e22 = h11(e14)) | (e22 = h11(e13)) | (e22 = h11(e12)) | (e22 = h11(e11)) | (e22 = h11(e10))) & ((e21 = h11(e15)) | (e21 = h11(e14)) | (e21 = h11(e13)) | (e21 = h11(e12)) | (e21 = h11(e11)) | (e21 = h11(e10))) & ((e20 = h11(e15)) | (e20 = h11(e14)) | (e20 = h11(e13)) | (e20 = h11(e12)) | (e20 = h11(e11)) | (e20 = h11(e10))) & (h11(op1(e15, e15)) = op2(h11(e15), h11(e15))) & (h11(op1(e15, e14)) = op2(h11(e15), h11(e14))) & (h11(op1(e15, e13)) = op2(h11(e15), h11(e13))) & (h11(op1(e15, e12)) = op2(h11(e15), h11(e12))) & (h11(op1(e15, e11)) = op2(h11(e15), h11(e11))) & (h11(op1(e15, e10)) = op2(h11(e15), h11(e10))) & (h11(op1(e14, e15)) = op2(h11(e14), h11(e15))) & (h11(op1(e14, e14)) = op2(h11(e14), h11(e14))) & (h11(op1(e14, e13)) = op2(h11(e14), h11(e13))) & (h11(op1(e14, e12)) = op2(h11(e14), h11(e12))) & (h11(op1(e14, e11)) = op2(h11(e14), h11(e11))) & (h11(op1(e14, e10)) = op2(h11(e14), h11(e10))) & (h11(op1(e13, e15)) = op2(h11(e13), h11(e15))) & (h11(op1(e13, e14)) = op2(h11(e13), h11(e14))) & (h11(op1(e13, e13)) = op2(h11(e13), h11(e13))) & (h11(op1(e13, e12)) = op2(h11(e13), h11(e12))) & (h11(op1(e13, e11)) = op2(h11(e13), h11(e11))) & (h11(op1(e13, e10)) = op2(h11(e13), h11(e10))) & (h11(op1(e12, e15)) = op2(h11(e12), h11(e15))) & (h11(op1(e12, e14)) = op2(h11(e12), h11(e14))) & (h11(op1(e12, e13)) = op2(h11(e12), h11(e13))) & (h11(op1(e12, e12)) = op2(h11(e12), h11(e12))) & (h11(op1(e12, e11)) = op2(h11(e12), h11(e11))) & (h11(op1(e12, e10)) = op2(h11(e12), h11(e10))) & (h11(op1(e11, e15)) = op2(h11(e11), h11(e15))) & (h11(op1(e11, e14)) = op2(h11(e11), h11(e14))) & (h11(op1(e11, e13)) = op2(h11(e11), h11(e13))) & (h11(op1(e11, e12)) = op2(h11(e11), h11(e12))) & (h11(op1(e11, e11)) = op2(h11(e11), h11(e11))) & (h11(op1(e11, e10)) = op2(h11(e11), h11(e10))) & (h11(op1(e10, e15)) = op2(h11(e10), h11(e15))) & (h11(op1(e10, e14)) = op2(h11(e10), h11(e14))) & (h11(op1(e10, e13)) = op2(h11(e10), h11(e13))) & (h11(op1(e10, e12)) = op2(h11(e10), h11(e12))) & (h11(op1(e10, e11)) = op2(h11(e10), h11(e11))) & (h11(op1(e10, e10)) = op2(h11(e10), h11(e10)))) | (((e25 = h10(e15)) | (e25 = h10(e14)) | (e25 = h10(e13)) | (e25 = h10(e12)) | (e25 = h10(e11)) | (e25 = h10(e10))) & ((e24 = h10(e15)) | (e24 = h10(e14)) | (e24 = h10(e13)) | (e24 = h10(e12)) | (e24 = h10(e11)) | (e24 = h10(e10))) & ((e23 = h10(e15)) | (e23 = h10(e14)) | (e23 = h10(e13)) | (e23 = h10(e12)) | (e23 = h10(e11)) | (e23 = h10(e10))) & ((e22 = h10(e15)) | (e22 = h10(e14)) | (e22 = h10(e13)) | (e22 = h10(e12)) | (e22 = h10(e11)) | (e22 = h10(e10))) & ((e21 = h10(e15)) | (e21 = h10(e14)) | (e21 = h10(e13)) | (e21 = h10(e12)) | (e21 = h10(e11)) | (e21 = h10(e10))) & ((e20 = h10(e15)) | (e20 = h10(e14)) | (e20 = h10(e13)) | (e20 = h10(e12)) | (e20 = h10(e11)) | (e20 = h10(e10))) & (h10(op1(e15, e15)) = op2(h10(e15), h10(e15))) & (h10(op1(e15, e14)) = op2(h10(e15), h10(e14))) & (h10(op1(e15, e13)) = op2(h10(e15), h10(e13))) & (h10(op1(e15, e12)) = op2(h10(e15), h10(e12))) & (h10(op1(e15, e11)) = op2(h10(e15), h10(e11))) & (h10(op1(e15, e10)) = op2(h10(e15), h10(e10))) & (h10(op1(e14, e15)) = op2(h10(e14), h10(e15))) & (h10(op1(e14, e14)) = op2(h10(e14), h10(e14))) & (h10(op1(e14, e13)) = op2(h10(e14), h10(e13))) & (h10(op1(e14, e12)) = op2(h10(e14), h10(e12))) & (h10(op1(e14, e11)) = op2(h10(e14), h10(e11))) & (h10(op1(e14, e10)) = op2(h10(e14), h10(e10))) & (h10(op1(e13, e15)) = op2(h10(e13), h10(e15))) & (h10(op1(e13, e14)) = op2(h10(e13), h10(e14))) & (h10(op1(e13, e13)) = op2(h10(e13), h10(e13))) & (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) & (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) & (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) & (h10(op1(e12, e15)) = op2(h10(e12), h10(e15))) & (h10(op1(e12, e14)) = op2(h10(e12), h10(e14))) & (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) & (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) & (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) & (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) & (h10(op1(e11, e15)) = op2(h10(e11), h10(e15))) & (h10(op1(e11, e14)) = op2(h10(e11), h10(e14))) & (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) & (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) & (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) & (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) & (h10(op1(e10, e15)) = op2(h10(e10), h10(e15))) & (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) & (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) & (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) & (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) & (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))) | (((e25 = h9(e15)) | (e25 = h9(e14)) | (e25 = h9(e13)) | (e25 = h9(e12)) | (e25 = h9(e11)) | (e25 = h9(e10))) & ((e24 = h9(e15)) | (e24 = h9(e14)) | (e24 = h9(e13)) | (e24 = h9(e12)) | (e24 = h9(e11)) | (e24 = h9(e10))) & ((e23 = h9(e15)) | (e23 = h9(e14)) | (e23 = h9(e13)) | (e23 = h9(e12)) | (e23 = h9(e11)) | (e23 = h9(e10))) & ((e22 = h9(e15)) | (e22 = h9(e14)) | (e22 = h9(e13)) | (e22 = h9(e12)) | (e22 = h9(e11)) | (e22 = h9(e10))) & ((e21 = h9(e15)) | (e21 = h9(e14)) | (e21 = h9(e13)) | (e21 = h9(e12)) | (e21 = h9(e11)) | (e21 = h9(e10))) & ((e20 = h9(e15)) | (e20 = h9(e14)) | (e20 = h9(e13)) | (e20 = h9(e12)) | (e20 = h9(e11)) | (e20 = h9(e10))) & (h9(op1(e15, e15)) = op2(h9(e15), h9(e15))) & (h9(op1(e15, e14)) = op2(h9(e15), h9(e14))) & (h9(op1(e15, e13)) = op2(h9(e15), h9(e13))) & (h9(op1(e15, e12)) = op2(h9(e15), h9(e12))) & (h9(op1(e15, e11)) = op2(h9(e15), h9(e11))) & (h9(op1(e15, e10)) = op2(h9(e15), h9(e10))) & (h9(op1(e14, e15)) = op2(h9(e14), h9(e15))) & (h9(op1(e14, e14)) = op2(h9(e14), h9(e14))) & (h9(op1(e14, e13)) = op2(h9(e14), h9(e13))) & (h9(op1(e14, e12)) = op2(h9(e14), h9(e12))) & (h9(op1(e14, e11)) = op2(h9(e14), h9(e11))) & (h9(op1(e14, e10)) = op2(h9(e14), h9(e10))) & (h9(op1(e13, e15)) = op2(h9(e13), h9(e15))) & (h9(op1(e13, e14)) = op2(h9(e13), h9(e14))) & (h9(op1(e13, e13)) = op2(h9(e13), h9(e13))) & (h9(op1(e13, e12)) = op2(h9(e13), h9(e12))) & (h9(op1(e13, e11)) = op2(h9(e13), h9(e11))) & (h9(op1(e13, e10)) = op2(h9(e13), h9(e10))) & (h9(op1(e12, e15)) = op2(h9(e12), h9(e15))) & (h9(op1(e12, e14)) = op2(h9(e12), h9(e14))) & (h9(op1(e12, e13)) = op2(h9(e12), h9(e13))) & (h9(op1(e12, e12)) = op2(h9(e12), h9(e12))) & (h9(op1(e12, e11)) = op2(h9(e12), h9(e11))) & (h9(op1(e12, e10)) = op2(h9(e12), h9(e10))) & (h9(op1(e11, e15)) = op2(h9(e11), h9(e15))) & (h9(op1(e11, e14)) = op2(h9(e11), h9(e14))) & (h9(op1(e11, e13)) = op2(h9(e11), h9(e13))) & (h9(op1(e11, e12)) = op2(h9(e11), h9(e12))) & (h9(op1(e11, e11)) = op2(h9(e11), h9(e11))) & (h9(op1(e11, e10)) = op2(h9(e11), h9(e10))) & (h9(op1(e10, e15)) = op2(h9(e10), h9(e15))) & (h9(op1(e10, e14)) = op2(h9(e10), h9(e14))) & (h9(op1(e10, e13)) = op2(h9(e10), h9(e13))) & (h9(op1(e10, e12)) = op2(h9(e10), h9(e12))) & (h9(op1(e10, e11)) = op2(h9(e10), h9(e11))) & (h9(op1(e10, e10)) = op2(h9(e10), h9(e10)))) | (((e25 = h8(e15)) | (e25 = h8(e14)) | (e25 = h8(e13)) | (e25 = h8(e12)) | (e25 = h8(e11)) | (e25 = h8(e10))) & ((e24 = h8(e15)) | (e24 = h8(e14)) | (e24 = h8(e13)) | (e24 = h8(e12)) | (e24 = h8(e11)) | (e24 = h8(e10))) & ((e23 = h8(e15)) | (e23 = h8(e14)) | (e23 = h8(e13)) | (e23 = h8(e12)) | (e23 = h8(e11)) | (e23 = h8(e10))) & ((e22 = h8(e15)) | (e22 = h8(e14)) | (e22 = h8(e13)) | (e22 = h8(e12)) | (e22 = h8(e11)) | (e22 = h8(e10))) & ((e21 = h8(e15)) | (e21 = h8(e14)) | (e21 = h8(e13)) | (e21 = h8(e12)) | (e21 = h8(e11)) | (e21 = h8(e10))) & ((e20 = h8(e15)) | (e20 = h8(e14)) | (e20 = h8(e13)) | (e20 = h8(e12)) | (e20 = h8(e11)) | (e20 = h8(e10))) & (h8(op1(e15, e15)) = op2(h8(e15), h8(e15))) & (h8(op1(e15, e14)) = op2(h8(e15), h8(e14))) & (h8(op1(e15, e13)) = op2(h8(e15), h8(e13))) & (h8(op1(e15, e12)) = op2(h8(e15), h8(e12))) & (h8(op1(e15, e11)) = op2(h8(e15), h8(e11))) & (h8(op1(e15, e10)) = op2(h8(e15), h8(e10))) & (h8(op1(e14, e15)) = op2(h8(e14), h8(e15))) & (h8(op1(e14, e14)) = op2(h8(e14), h8(e14))) & (h8(op1(e14, e13)) = op2(h8(e14), h8(e13))) & (h8(op1(e14, e12)) = op2(h8(e14), h8(e12))) & (h8(op1(e14, e11)) = op2(h8(e14), h8(e11))) & (h8(op1(e14, e10)) = op2(h8(e14), h8(e10))) & (h8(op1(e13, e15)) = op2(h8(e13), h8(e15))) & (h8(op1(e13, e14)) = op2(h8(e13), h8(e14))) & (h8(op1(e13, e13)) = op2(h8(e13), h8(e13))) & (h8(op1(e13, e12)) = op2(h8(e13), h8(e12))) & (h8(op1(e13, e11)) = op2(h8(e13), h8(e11))) & (h8(op1(e13, e10)) = op2(h8(e13), h8(e10))) & (h8(op1(e12, e15)) = op2(h8(e12), h8(e15))) & (h8(op1(e12, e14)) = op2(h8(e12), h8(e14))) & (h8(op1(e12, e13)) = op2(h8(e12), h8(e13))) & (h8(op1(e12, e12)) = op2(h8(e12), h8(e12))) & (h8(op1(e12, e11)) = op2(h8(e12), h8(e11))) & (h8(op1(e12, e10)) = op2(h8(e12), h8(e10))) & (h8(op1(e11, e15)) = op2(h8(e11), h8(e15))) & (h8(op1(e11, e14)) = op2(h8(e11), h8(e14))) & (h8(op1(e11, e13)) = op2(h8(e11), h8(e13))) & (h8(op1(e11, e12)) = op2(h8(e11), h8(e12))) & (h8(op1(e11, e11)) = op2(h8(e11), h8(e11))) & (h8(op1(e11, e10)) = op2(h8(e11), h8(e10))) & (h8(op1(e10, e15)) = op2(h8(e10), h8(e15))) & (h8(op1(e10, e14)) = op2(h8(e10), h8(e14))) & (h8(op1(e10, e13)) = op2(h8(e10), h8(e13))) & (h8(op1(e10, e12)) = op2(h8(e10), h8(e12))) & (h8(op1(e10, e11)) = op2(h8(e10), h8(e11))) & (h8(op1(e10, e10)) = op2(h8(e10), h8(e10)))) | (((e25 = h7(e15)) | (e25 = h7(e14)) | (e25 = h7(e13)) | (e25 = h7(e12)) | (e25 = h7(e11)) | (e25 = h7(e10))) & ((e24 = h7(e15)) | (e24 = h7(e14)) | (e24 = h7(e13)) | (e24 = h7(e12)) | (e24 = h7(e11)) | (e24 = h7(e10))) & ((e23 = h7(e15)) | (e23 = h7(e14)) | (e23 = h7(e13)) | (e23 = h7(e12)) | (e23 = h7(e11)) | (e23 = h7(e10))) & ((e22 = h7(e15)) | (e22 = h7(e14)) | (e22 = h7(e13)) | (e22 = h7(e12)) | (e22 = h7(e11)) | (e22 = h7(e10))) & ((e21 = h7(e15)) | (e21 = h7(e14)) | (e21 = h7(e13)) | (e21 = h7(e12)) | (e21 = h7(e11)) | (e21 = h7(e10))) & ((e20 = h7(e15)) | (e20 = h7(e14)) | (e20 = h7(e13)) | (e20 = h7(e12)) | (e20 = h7(e11)) | (e20 = h7(e10))) & (h7(op1(e15, e15)) = op2(h7(e15), h7(e15))) & (h7(op1(e15, e14)) = op2(h7(e15), h7(e14))) & (h7(op1(e15, e13)) = op2(h7(e15), h7(e13))) & (h7(op1(e15, e12)) = op2(h7(e15), h7(e12))) & (h7(op1(e15, e11)) = op2(h7(e15), h7(e11))) & (h7(op1(e15, e10)) = op2(h7(e15), h7(e10))) & (h7(op1(e14, e15)) = op2(h7(e14), h7(e15))) & (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) & (h7(op1(e14, e13)) = op2(h7(e14), h7(e13))) & (h7(op1(e14, e12)) = op2(h7(e14), h7(e12))) & (h7(op1(e14, e11)) = op2(h7(e14), h7(e11))) & (h7(op1(e14, e10)) = op2(h7(e14), h7(e10))) & (h7(op1(e13, e15)) = op2(h7(e13), h7(e15))) & (h7(op1(e13, e14)) = op2(h7(e13), h7(e14))) & (h7(op1(e13, e13)) = op2(h7(e13), h7(e13))) & (h7(op1(e13, e12)) = op2(h7(e13), h7(e12))) & (h7(op1(e13, e11)) = op2(h7(e13), h7(e11))) & (h7(op1(e13, e10)) = op2(h7(e13), h7(e10))) & (h7(op1(e12, e15)) = op2(h7(e12), h7(e15))) & (h7(op1(e12, e14)) = op2(h7(e12), h7(e14))) & (h7(op1(e12, e13)) = op2(h7(e12), h7(e13))) & (h7(op1(e12, e12)) = op2(h7(e12), h7(e12))) & (h7(op1(e12, e11)) = op2(h7(e12), h7(e11))) & (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) & (h7(op1(e11, e15)) = op2(h7(e11), h7(e15))) & (h7(op1(e11, e14)) = op2(h7(e11), h7(e14))) & (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) & (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) & (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) & (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) & (h7(op1(e10, e15)) = op2(h7(e10), h7(e15))) & (h7(op1(e10, e14)) = op2(h7(e10), h7(e14))) & (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) & (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) & (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) & (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))) | (((e25 = h6(e15)) | (e25 = h6(e14)) | (e25 = h6(e13)) | (e25 = h6(e12)) | (e25 = h6(e11)) | (e25 = h6(e10))) & ((e24 = h6(e15)) | (e24 = h6(e14)) | (e24 = h6(e13)) | (e24 = h6(e12)) | (e24 = h6(e11)) | (e24 = h6(e10))) & ((e23 = h6(e15)) | (e23 = h6(e14)) | (e23 = h6(e13)) | (e23 = h6(e12)) | (e23 = h6(e11)) | (e23 = h6(e10))) & ((e22 = h6(e15)) | (e22 = h6(e14)) | (e22 = h6(e13)) | (e22 = h6(e12)) | (e22 = h6(e11)) | (e22 = h6(e10))) & ((e21 = h6(e15)) | (e21 = h6(e14)) | (e21 = h6(e13)) | (e21 = h6(e12)) | (e21 = h6(e11)) | (e21 = h6(e10))) & ((e20 = h6(e15)) | (e20 = h6(e14)) | (e20 = h6(e13)) | (e20 = h6(e12)) | (e20 = h6(e11)) | (e20 = h6(e10))) & (h6(op1(e15, e15)) = op2(h6(e15), h6(e15))) & (h6(op1(e15, e14)) = op2(h6(e15), h6(e14))) & (h6(op1(e15, e13)) = op2(h6(e15), h6(e13))) & (h6(op1(e15, e12)) = op2(h6(e15), h6(e12))) & (h6(op1(e15, e11)) = op2(h6(e15), h6(e11))) & (h6(op1(e15, e10)) = op2(h6(e15), h6(e10))) & (h6(op1(e14, e15)) = op2(h6(e14), h6(e15))) & (h6(op1(e14, e14)) = op2(h6(e14), h6(e14))) & (h6(op1(e14, e13)) = op2(h6(e14), h6(e13))) & (h6(op1(e14, e12)) = op2(h6(e14), h6(e12))) & (h6(op1(e14, e11)) = op2(h6(e14), h6(e11))) & (h6(op1(e14, e10)) = op2(h6(e14), h6(e10))) & (h6(op1(e13, e15)) = op2(h6(e13), h6(e15))) & (h6(op1(e13, e14)) = op2(h6(e13), h6(e14))) & (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) & (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) & (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) & (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) & (h6(op1(e12, e15)) = op2(h6(e12), h6(e15))) & (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) & (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) & (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) & (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) & (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) & (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) & (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) & (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) & (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) & (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) & (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) & (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) & (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) & (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) & (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) & (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) & (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))) | (((e25 = h5(e15)) | (e25 = h5(e14)) | (e25 = h5(e13)) | (e25 = h5(e12)) | (e25 = h5(e11)) | (e25 = h5(e10))) & ((e24 = h5(e15)) | (e24 = h5(e14)) | (e24 = h5(e13)) | (e24 = h5(e12)) | (e24 = h5(e11)) | (e24 = h5(e10))) & ((e23 = h5(e15)) | (e23 = h5(e14)) | (e23 = h5(e13)) | (e23 = h5(e12)) | (e23 = h5(e11)) | (e23 = h5(e10))) & ((e22 = h5(e15)) | (e22 = h5(e14)) | (e22 = h5(e13)) | (e22 = h5(e12)) | (e22 = h5(e11)) | (e22 = h5(e10))) & ((e21 = h5(e15)) | (e21 = h5(e14)) | (e21 = h5(e13)) | (e21 = h5(e12)) | (e21 = h5(e11)) | (e21 = h5(e10))) & ((e20 = h5(e15)) | (e20 = h5(e14)) | (e20 = h5(e13)) | (e20 = h5(e12)) | (e20 = h5(e11)) | (e20 = h5(e10))) & (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) & (h5(op1(e15, e14)) = op2(h5(e15), h5(e14))) & (h5(op1(e15, e13)) = op2(h5(e15), h5(e13))) & (h5(op1(e15, e12)) = op2(h5(e15), h5(e12))) & (h5(op1(e15, e11)) = op2(h5(e15), h5(e11))) & (h5(op1(e15, e10)) = op2(h5(e15), h5(e10))) & (h5(op1(e14, e15)) = op2(h5(e14), h5(e15))) & (h5(op1(e14, e14)) = op2(h5(e14), h5(e14))) & (h5(op1(e14, e13)) = op2(h5(e14), h5(e13))) & (h5(op1(e14, e12)) = op2(h5(e14), h5(e12))) & (h5(op1(e14, e11)) = op2(h5(e14), h5(e11))) & (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) & (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) & (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) & (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) & (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) & (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) & (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) & (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) & (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) & (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) & (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) & (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) & (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) & (h5(op1(e11, e15)) = op2(h5(e11), h5(e15))) & (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) & (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) & (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) & (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) & (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) & (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) & (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) & (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) & (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) & (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) & (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))) | (((e25 = h4(e15)) | (e25 = h4(e14)) | (e25 = h4(e13)) | (e25 = h4(e12)) | (e25 = h4(e11)) | (e25 = h4(e10))) & ((e24 = h4(e15)) | (e24 = h4(e14)) | (e24 = h4(e13)) | (e24 = h4(e12)) | (e24 = h4(e11)) | (e24 = h4(e10))) & ((e23 = h4(e15)) | (e23 = h4(e14)) | (e23 = h4(e13)) | (e23 = h4(e12)) | (e23 = h4(e11)) | (e23 = h4(e10))) & ((e22 = h4(e15)) | (e22 = h4(e14)) | (e22 = h4(e13)) | (e22 = h4(e12)) | (e22 = h4(e11)) | (e22 = h4(e10))) & ((e21 = h4(e15)) | (e21 = h4(e14)) | (e21 = h4(e13)) | (e21 = h4(e12)) | (e21 = h4(e11)) | (e21 = h4(e10))) & ((e20 = h4(e15)) | (e20 = h4(e14)) | (e20 = h4(e13)) | (e20 = h4(e12)) | (e20 = h4(e11)) | (e20 = h4(e10))) & (h4(op1(e15, e15)) = op2(h4(e15), h4(e15))) & (h4(op1(e15, e14)) = op2(h4(e15), h4(e14))) & (h4(op1(e15, e13)) = op2(h4(e15), h4(e13))) & (h4(op1(e15, e12)) = op2(h4(e15), h4(e12))) & (h4(op1(e15, e11)) = op2(h4(e15), h4(e11))) & (h4(op1(e15, e10)) = op2(h4(e15), h4(e10))) & (h4(op1(e14, e15)) = op2(h4(e14), h4(e15))) & (h4(op1(e14, e14)) = op2(h4(e14), h4(e14))) & (h4(op1(e14, e13)) = op2(h4(e14), h4(e13))) & (h4(op1(e14, e12)) = op2(h4(e14), h4(e12))) & (h4(op1(e14, e11)) = op2(h4(e14), h4(e11))) & (h4(op1(e14, e10)) = op2(h4(e14), h4(e10))) & (h4(op1(e13, e15)) = op2(h4(e13), h4(e15))) & (h4(op1(e13, e14)) = op2(h4(e13), h4(e14))) & (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) & (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) & (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) & (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) & (h4(op1(e12, e15)) = op2(h4(e12), h4(e15))) & (h4(op1(e12, e14)) = op2(h4(e12), h4(e14))) & (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) & (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) & (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) & (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) & (h4(op1(e11, e15)) = op2(h4(e11), h4(e15))) & (h4(op1(e11, e14)) = op2(h4(e11), h4(e14))) & (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) & (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) & (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) & (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) & (h4(op1(e10, e15)) = op2(h4(e10), h4(e15))) & (h4(op1(e10, e14)) = op2(h4(e10), h4(e14))) & (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) & (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) & (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) & (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) | (((e25 = h3(e15)) | (e25 = h3(e14)) | (e25 = h3(e13)) | (e25 = h3(e12)) | (e25 = h3(e11)) | (e25 = h3(e10))) & ((e24 = h3(e15)) | (e24 = h3(e14)) | (e24 = h3(e13)) | (e24 = h3(e12)) | (e24 = h3(e11)) | (e24 = h3(e10))) & ((e23 = h3(e15)) | (e23 = h3(e14)) | (e23 = h3(e13)) | (e23 = h3(e12)) | (e23 = h3(e11)) | (e23 = h3(e10))) & ((e22 = h3(e15)) | (e22 = h3(e14)) | (e22 = h3(e13)) | (e22 = h3(e12)) | (e22 = h3(e11)) | (e22 = h3(e10))) & ((e21 = h3(e15)) | (e21 = h3(e14)) | (e21 = h3(e13)) | (e21 = h3(e12)) | (e21 = h3(e11)) | (e21 = h3(e10))) & ((e20 = h3(e15)) | (e20 = h3(e14)) | (e20 = h3(e13)) | (e20 = h3(e12)) | (e20 = h3(e11)) | (e20 = h3(e10))) & (h3(op1(e15, e15)) = op2(h3(e15), h3(e15))) & (h3(op1(e15, e14)) = op2(h3(e15), h3(e14))) & (h3(op1(e15, e13)) = op2(h3(e15), h3(e13))) & (h3(op1(e15, e12)) = op2(h3(e15), h3(e12))) & (h3(op1(e15, e11)) = op2(h3(e15), h3(e11))) & (h3(op1(e15, e10)) = op2(h3(e15), h3(e10))) & (h3(op1(e14, e15)) = op2(h3(e14), h3(e15))) & (h3(op1(e14, e14)) = op2(h3(e14), h3(e14))) & (h3(op1(e14, e13)) = op2(h3(e14), h3(e13))) & (h3(op1(e14, e12)) = op2(h3(e14), h3(e12))) & (h3(op1(e14, e11)) = op2(h3(e14), h3(e11))) & (h3(op1(e14, e10)) = op2(h3(e14), h3(e10))) & (h3(op1(e13, e15)) = op2(h3(e13), h3(e15))) & (h3(op1(e13, e14)) = op2(h3(e13), h3(e14))) & (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) & (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) & (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) & (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) & (h3(op1(e12, e15)) = op2(h3(e12), h3(e15))) & (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) & (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) & (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) & (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) & (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) & (h3(op1(e11, e15)) = op2(h3(e11), h3(e15))) & (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) & (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) & (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) & (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) & (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) & (h3(op1(e10, e15)) = op2(h3(e10), h3(e15))) & (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) & (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) & (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) & (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) & (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) | (((e25 = h2(e15)) | (e25 = h2(e14)) | (e25 = h2(e13)) | (e25 = h2(e12)) | (e25 = h2(e11)) | (e25 = h2(e10))) & ((e24 = h2(e15)) | (e24 = h2(e14)) | (e24 = h2(e13)) | (e24 = h2(e12)) | (e24 = h2(e11)) | (e24 = h2(e10))) & ((e23 = h2(e15)) | (e23 = h2(e14)) | (e23 = h2(e13)) | (e23 = h2(e12)) | (e23 = h2(e11)) | (e23 = h2(e10))) & ((e22 = h2(e15)) | (e22 = h2(e14)) | (e22 = h2(e13)) | (e22 = h2(e12)) | (e22 = h2(e11)) | (e22 = h2(e10))) & ((e21 = h2(e15)) | (e21 = h2(e14)) | (e21 = h2(e13)) | (e21 = h2(e12)) | (e21 = h2(e11)) | (e21 = h2(e10))) & ((e20 = h2(e15)) | (e20 = h2(e14)) | (e20 = h2(e13)) | (e20 = h2(e12)) | (e20 = h2(e11)) | (e20 = h2(e10))) & (h2(op1(e15, e15)) = op2(h2(e15), h2(e15))) & (h2(op1(e15, e14)) = op2(h2(e15), h2(e14))) & (h2(op1(e15, e13)) = op2(h2(e15), h2(e13))) & (h2(op1(e15, e12)) = op2(h2(e15), h2(e12))) & (h2(op1(e15, e11)) = op2(h2(e15), h2(e11))) & (h2(op1(e15, e10)) = op2(h2(e15), h2(e10))) & (h2(op1(e14, e15)) = op2(h2(e14), h2(e15))) & (h2(op1(e14, e14)) = op2(h2(e14), h2(e14))) & (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) & (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) & (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) & (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) & (h2(op1(e13, e15)) = op2(h2(e13), h2(e15))) & (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) & (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) & (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) & (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) & (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) & (h2(op1(e12, e15)) = op2(h2(e12), h2(e15))) & (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) & (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) & (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) & (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) & (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) & (h2(op1(e11, e15)) = op2(h2(e11), h2(e15))) & (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) & (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) & (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) & (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) & (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) & (h2(op1(e10, e15)) = op2(h2(e10), h2(e15))) & (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) & (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) & (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) & (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) & (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) | (((e25 = h1(e15)) | (e25 = h1(e14)) | (e25 = h1(e13)) | (e25 = h1(e12)) | (e25 = h1(e11)) | (e25 = h1(e10))) & ((e24 = h1(e15)) | (e24 = h1(e14)) | (e24 = h1(e13)) | (e24 = h1(e12)) | (e24 = h1(e11)) | (e24 = h1(e10))) & ((e23 = h1(e15)) | (e23 = h1(e14)) | (e23 = h1(e13)) | (e23 = h1(e12)) | (e23 = h1(e11)) | (e23 = h1(e10))) & ((e22 = h1(e15)) | (e22 = h1(e14)) | (e22 = h1(e13)) | (e22 = h1(e12)) | (e22 = h1(e11)) | (e22 = h1(e10))) & ((e21 = h1(e15)) | (e21 = h1(e14)) | (e21 = h1(e13)) | (e21 = h1(e12)) | (e21 = h1(e11)) | (e21 = h1(e10))) & ((e20 = h1(e15)) | (e20 = h1(e14)) | (e20 = h1(e13)) | (e20 = h1(e12)) | (e20 = h1(e11)) | (e20 = h1(e10))) & (h1(op1(e15, e15)) = op2(h1(e15), h1(e15))) & (h1(op1(e15, e14)) = op2(h1(e15), h1(e14))) & (h1(op1(e15, e13)) = op2(h1(e15), h1(e13))) & (h1(op1(e15, e12)) = op2(h1(e15), h1(e12))) & (h1(op1(e15, e11)) = op2(h1(e15), h1(e11))) & (h1(op1(e15, e10)) = op2(h1(e15), h1(e10))) & (h1(op1(e14, e15)) = op2(h1(e14), h1(e15))) & (h1(op1(e14, e14)) = op2(h1(e14), h1(e14))) & (h1(op1(e14, e13)) = op2(h1(e14), h1(e13))) & (h1(op1(e14, e12)) = op2(h1(e14), h1(e12))) & (h1(op1(e14, e11)) = op2(h1(e14), h1(e11))) & (h1(op1(e14, e10)) = op2(h1(e14), h1(e10))) & (h1(op1(e13, e15)) = op2(h1(e13), h1(e15))) & (h1(op1(e13, e14)) = op2(h1(e13), h1(e14))) & (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) & (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) & (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) & (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) & (h1(op1(e12, e15)) = op2(h1(e12), h1(e15))) & (h1(op1(e12, e14)) = op2(h1(e12), h1(e14))) & (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) & (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) & (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) & (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) & (h1(op1(e11, e15)) = op2(h1(e11), h1(e15))) & (h1(op1(e11, e14)) = op2(h1(e11), h1(e14))) & (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) & (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) & (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) & (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) & (h1(op1(e10, e15)) = op2(h1(e10), h1(e15))) & (h1(op1(e10, e14)) = op2(h1(e10), h1(e14))) & (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) & (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) & (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) & (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(negated_conjecture, [], [f56])).
fof(f56, plain, ~ ((((e25 = h30(e15)) | (e25 = h30(e14)) | (e25 = h30(e13)) | (e25 = h30(e12)) | (e25 = h30(e11)) | (e25 = h30(e10))) & ((e24 = h30(e15)) | (e24 = h30(e14)) | (e24 = h30(e13)) | (e24 = h30(e12)) | (e24 = h30(e11)) | (e24 = h30(e10))) & ((e23 = h30(e15)) | (e23 = h30(e14)) | (e23 = h30(e13)) | (e23 = h30(e12)) | (e23 = h30(e11)) | (e23 = h30(e10))) & ((e22 = h30(e15)) | (e22 = h30(e14)) | (e22 = h30(e13)) | (e22 = h30(e12)) | (e22 = h30(e11)) | (e22 = h30(e10))) & ((e21 = h30(e15)) | (e21 = h30(e14)) | (e21 = h30(e13)) | (e21 = h30(e12)) | (e21 = h30(e11)) | (e21 = h30(e10))) & ((e20 = h30(e15)) | (e20 = h30(e14)) | (e20 = h30(e13)) | (e20 = h30(e12)) | (e20 = h30(e11)) | (e20 = h30(e10))) & (h30(op1(e15, e15)) = op2(h30(e15), h30(e15))) & (h30(op1(e15, e14)) = op2(h30(e15), h30(e14))) & (h30(op1(e15, e13)) = op2(h30(e15), h30(e13))) & (h30(op1(e15, e12)) = op2(h30(e15), h30(e12))) & (h30(op1(e15, e11)) = op2(h30(e15), h30(e11))) & (h30(op1(e15, e10)) = op2(h30(e15), h30(e10))) & (h30(op1(e14, e15)) = op2(h30(e14), h30(e15))) & (h30(op1(e14, e14)) = op2(h30(e14), h30(e14))) & (h30(op1(e14, e13)) = op2(h30(e14), h30(e13))) & (h30(op1(e14, e12)) = op2(h30(e14), h30(e12))) & (h30(op1(e14, e11)) = op2(h30(e14), h30(e11))) & (h30(op1(e14, e10)) = op2(h30(e14), h30(e10))) & (h30(op1(e13, e15)) = op2(h30(e13), h30(e15))) & (h30(op1(e13, e14)) = op2(h30(e13), h30(e14))) & (h30(op1(e13, e13)) = op2(h30(e13), h30(e13))) & (h30(op1(e13, e12)) = op2(h30(e13), h30(e12))) & (h30(op1(e13, e11)) = op2(h30(e13), h30(e11))) & (h30(op1(e13, e10)) = op2(h30(e13), h30(e10))) & (h30(op1(e12, e15)) = op2(h30(e12), h30(e15))) & (h30(op1(e12, e14)) = op2(h30(e12), h30(e14))) & (h30(op1(e12, e13)) = op2(h30(e12), h30(e13))) & (h30(op1(e12, e12)) = op2(h30(e12), h30(e12))) & (h30(op1(e12, e11)) = op2(h30(e12), h30(e11))) & (h30(op1(e12, e10)) = op2(h30(e12), h30(e10))) & (h30(op1(e11, e15)) = op2(h30(e11), h30(e15))) & (h30(op1(e11, e14)) = op2(h30(e11), h30(e14))) & (h30(op1(e11, e13)) = op2(h30(e11), h30(e13))) & (h30(op1(e11, e12)) = op2(h30(e11), h30(e12))) & (h30(op1(e11, e11)) = op2(h30(e11), h30(e11))) & (h30(op1(e11, e10)) = op2(h30(e11), h30(e10))) & (h30(op1(e10, e15)) = op2(h30(e10), h30(e15))) & (h30(op1(e10, e14)) = op2(h30(e10), h30(e14))) & (h30(op1(e10, e13)) = op2(h30(e10), h30(e13))) & (h30(op1(e10, e12)) = op2(h30(e10), h30(e12))) & (h30(op1(e10, e11)) = op2(h30(e10), h30(e11))) & (h30(op1(e10, e10)) = op2(h30(e10), h30(e10)))) | (((e25 = h29(e15)) | (e25 = h29(e14)) | (e25 = h29(e13)) | (e25 = h29(e12)) | (e25 = h29(e11)) | (e25 = h29(e10))) & ((e24 = h29(e15)) | (e24 = h29(e14)) | (e24 = h29(e13)) | (e24 = h29(e12)) | (e24 = h29(e11)) | (e24 = h29(e10))) & ((e23 = h29(e15)) | (e23 = h29(e14)) | (e23 = h29(e13)) | (e23 = h29(e12)) | (e23 = h29(e11)) | (e23 = h29(e10))) & ((e22 = h29(e15)) | (e22 = h29(e14)) | (e22 = h29(e13)) | (e22 = h29(e12)) | (e22 = h29(e11)) | (e22 = h29(e10))) & ((e21 = h29(e15)) | (e21 = h29(e14)) | (e21 = h29(e13)) | (e21 = h29(e12)) | (e21 = h29(e11)) | (e21 = h29(e10))) & ((e20 = h29(e15)) | (e20 = h29(e14)) | (e20 = h29(e13)) | (e20 = h29(e12)) | (e20 = h29(e11)) | (e20 = h29(e10))) & (h29(op1(e15, e15)) = op2(h29(e15), h29(e15))) & (h29(op1(e15, e14)) = op2(h29(e15), h29(e14))) & (h29(op1(e15, e13)) = op2(h29(e15), h29(e13))) & (h29(op1(e15, e12)) = op2(h29(e15), h29(e12))) & (h29(op1(e15, e11)) = op2(h29(e15), h29(e11))) & (h29(op1(e15, e10)) = op2(h29(e15), h29(e10))) & (h29(op1(e14, e15)) = op2(h29(e14), h29(e15))) & (h29(op1(e14, e14)) = op2(h29(e14), h29(e14))) & (h29(op1(e14, e13)) = op2(h29(e14), h29(e13))) & (h29(op1(e14, e12)) = op2(h29(e14), h29(e12))) & (h29(op1(e14, e11)) = op2(h29(e14), h29(e11))) & (h29(op1(e14, e10)) = op2(h29(e14), h29(e10))) & (h29(op1(e13, e15)) = op2(h29(e13), h29(e15))) & (h29(op1(e13, e14)) = op2(h29(e13), h29(e14))) & (h29(op1(e13, e13)) = op2(h29(e13), h29(e13))) & (h29(op1(e13, e12)) = op2(h29(e13), h29(e12))) & (h29(op1(e13, e11)) = op2(h29(e13), h29(e11))) & (h29(op1(e13, e10)) = op2(h29(e13), h29(e10))) & (h29(op1(e12, e15)) = op2(h29(e12), h29(e15))) & (h29(op1(e12, e14)) = op2(h29(e12), h29(e14))) & (h29(op1(e12, e13)) = op2(h29(e12), h29(e13))) & (h29(op1(e12, e12)) = op2(h29(e12), h29(e12))) & (h29(op1(e12, e11)) = op2(h29(e12), h29(e11))) & (h29(op1(e12, e10)) = op2(h29(e12), h29(e10))) & (h29(op1(e11, e15)) = op2(h29(e11), h29(e15))) & (h29(op1(e11, e14)) = op2(h29(e11), h29(e14))) & (h29(op1(e11, e13)) = op2(h29(e11), h29(e13))) & (h29(op1(e11, e12)) = op2(h29(e11), h29(e12))) & (h29(op1(e11, e11)) = op2(h29(e11), h29(e11))) & (h29(op1(e11, e10)) = op2(h29(e11), h29(e10))) & (h29(op1(e10, e15)) = op2(h29(e10), h29(e15))) & (h29(op1(e10, e14)) = op2(h29(e10), h29(e14))) & (h29(op1(e10, e13)) = op2(h29(e10), h29(e13))) & (h29(op1(e10, e12)) = op2(h29(e10), h29(e12))) & (h29(op1(e10, e11)) = op2(h29(e10), h29(e11))) & (h29(op1(e10, e10)) = op2(h29(e10), h29(e10)))) | (((e25 = h28(e15)) | (e25 = h28(e14)) | (e25 = h28(e13)) | (e25 = h28(e12)) | (e25 = h28(e11)) | (e25 = h28(e10))) & ((e24 = h28(e15)) | (e24 = h28(e14)) | (e24 = h28(e13)) | (e24 = h28(e12)) | (e24 = h28(e11)) | (e24 = h28(e10))) & ((e23 = h28(e15)) | (e23 = h28(e14)) | (e23 = h28(e13)) | (e23 = h28(e12)) | (e23 = h28(e11)) | (e23 = h28(e10))) & ((e22 = h28(e15)) | (e22 = h28(e14)) | (e22 = h28(e13)) | (e22 = h28(e12)) | (e22 = h28(e11)) | (e22 = h28(e10))) & ((e21 = h28(e15)) | (e21 = h28(e14)) | (e21 = h28(e13)) | (e21 = h28(e12)) | (e21 = h28(e11)) | (e21 = h28(e10))) & ((e20 = h28(e15)) | (e20 = h28(e14)) | (e20 = h28(e13)) | (e20 = h28(e12)) | (e20 = h28(e11)) | (e20 = h28(e10))) & (h28(op1(e15, e15)) = op2(h28(e15), h28(e15))) & (h28(op1(e15, e14)) = op2(h28(e15), h28(e14))) & (h28(op1(e15, e13)) = op2(h28(e15), h28(e13))) & (h28(op1(e15, e12)) = op2(h28(e15), h28(e12))) & (h28(op1(e15, e11)) = op2(h28(e15), h28(e11))) & (h28(op1(e15, e10)) = op2(h28(e15), h28(e10))) & (h28(op1(e14, e15)) = op2(h28(e14), h28(e15))) & (h28(op1(e14, e14)) = op2(h28(e14), h28(e14))) & (h28(op1(e14, e13)) = op2(h28(e14), h28(e13))) & (h28(op1(e14, e12)) = op2(h28(e14), h28(e12))) & (h28(op1(e14, e11)) = op2(h28(e14), h28(e11))) & (h28(op1(e14, e10)) = op2(h28(e14), h28(e10))) & (h28(op1(e13, e15)) = op2(h28(e13), h28(e15))) & (h28(op1(e13, e14)) = op2(h28(e13), h28(e14))) & (h28(op1(e13, e13)) = op2(h28(e13), h28(e13))) & (h28(op1(e13, e12)) = op2(h28(e13), h28(e12))) & (h28(op1(e13, e11)) = op2(h28(e13), h28(e11))) & (h28(op1(e13, e10)) = op2(h28(e13), h28(e10))) & (h28(op1(e12, e15)) = op2(h28(e12), h28(e15))) & (h28(op1(e12, e14)) = op2(h28(e12), h28(e14))) & (h28(op1(e12, e13)) = op2(h28(e12), h28(e13))) & (h28(op1(e12, e12)) = op2(h28(e12), h28(e12))) & (h28(op1(e12, e11)) = op2(h28(e12), h28(e11))) & (h28(op1(e12, e10)) = op2(h28(e12), h28(e10))) & (h28(op1(e11, e15)) = op2(h28(e11), h28(e15))) & (h28(op1(e11, e14)) = op2(h28(e11), h28(e14))) & (h28(op1(e11, e13)) = op2(h28(e11), h28(e13))) & (h28(op1(e11, e12)) = op2(h28(e11), h28(e12))) & (h28(op1(e11, e11)) = op2(h28(e11), h28(e11))) & (h28(op1(e11, e10)) = op2(h28(e11), h28(e10))) & (h28(op1(e10, e15)) = op2(h28(e10), h28(e15))) & (h28(op1(e10, e14)) = op2(h28(e10), h28(e14))) & (h28(op1(e10, e13)) = op2(h28(e10), h28(e13))) & (h28(op1(e10, e12)) = op2(h28(e10), h28(e12))) & (h28(op1(e10, e11)) = op2(h28(e10), h28(e11))) & (h28(op1(e10, e10)) = op2(h28(e10), h28(e10)))) | (((e25 = h27(e15)) | (e25 = h27(e14)) | (e25 = h27(e13)) | (e25 = h27(e12)) | (e25 = h27(e11)) | (e25 = h27(e10))) & ((e24 = h27(e15)) | (e24 = h27(e14)) | (e24 = h27(e13)) | (e24 = h27(e12)) | (e24 = h27(e11)) | (e24 = h27(e10))) & ((e23 = h27(e15)) | (e23 = h27(e14)) | (e23 = h27(e13)) | (e23 = h27(e12)) | (e23 = h27(e11)) | (e23 = h27(e10))) & ((e22 = h27(e15)) | (e22 = h27(e14)) | (e22 = h27(e13)) | (e22 = h27(e12)) | (e22 = h27(e11)) | (e22 = h27(e10))) & ((e21 = h27(e15)) | (e21 = h27(e14)) | (e21 = h27(e13)) | (e21 = h27(e12)) | (e21 = h27(e11)) | (e21 = h27(e10))) & ((e20 = h27(e15)) | (e20 = h27(e14)) | (e20 = h27(e13)) | (e20 = h27(e12)) | (e20 = h27(e11)) | (e20 = h27(e10))) & (h27(op1(e15, e15)) = op2(h27(e15), h27(e15))) & (h27(op1(e15, e14)) = op2(h27(e15), h27(e14))) & (h27(op1(e15, e13)) = op2(h27(e15), h27(e13))) & (h27(op1(e15, e12)) = op2(h27(e15), h27(e12))) & (h27(op1(e15, e11)) = op2(h27(e15), h27(e11))) & (h27(op1(e15, e10)) = op2(h27(e15), h27(e10))) & (h27(op1(e14, e15)) = op2(h27(e14), h27(e15))) & (h27(op1(e14, e14)) = op2(h27(e14), h27(e14))) & (h27(op1(e14, e13)) = op2(h27(e14), h27(e13))) & (h27(op1(e14, e12)) = op2(h27(e14), h27(e12))) & (h27(op1(e14, e11)) = op2(h27(e14), h27(e11))) & (h27(op1(e14, e10)) = op2(h27(e14), h27(e10))) & (h27(op1(e13, e15)) = op2(h27(e13), h27(e15))) & (h27(op1(e13, e14)) = op2(h27(e13), h27(e14))) & (h27(op1(e13, e13)) = op2(h27(e13), h27(e13))) & (h27(op1(e13, e12)) = op2(h27(e13), h27(e12))) & (h27(op1(e13, e11)) = op2(h27(e13), h27(e11))) & (h27(op1(e13, e10)) = op2(h27(e13), h27(e10))) & (h27(op1(e12, e15)) = op2(h27(e12), h27(e15))) & (h27(op1(e12, e14)) = op2(h27(e12), h27(e14))) & (h27(op1(e12, e13)) = op2(h27(e12), h27(e13))) & (h27(op1(e12, e12)) = op2(h27(e12), h27(e12))) & (h27(op1(e12, e11)) = op2(h27(e12), h27(e11))) & (h27(op1(e12, e10)) = op2(h27(e12), h27(e10))) & (h27(op1(e11, e15)) = op2(h27(e11), h27(e15))) & (h27(op1(e11, e14)) = op2(h27(e11), h27(e14))) & (h27(op1(e11, e13)) = op2(h27(e11), h27(e13))) & (h27(op1(e11, e12)) = op2(h27(e11), h27(e12))) & (h27(op1(e11, e11)) = op2(h27(e11), h27(e11))) & (h27(op1(e11, e10)) = op2(h27(e11), h27(e10))) & (h27(op1(e10, e15)) = op2(h27(e10), h27(e15))) & (h27(op1(e10, e14)) = op2(h27(e10), h27(e14))) & (h27(op1(e10, e13)) = op2(h27(e10), h27(e13))) & (h27(op1(e10, e12)) = op2(h27(e10), h27(e12))) & (h27(op1(e10, e11)) = op2(h27(e10), h27(e11))) & (h27(op1(e10, e10)) = op2(h27(e10), h27(e10)))) | (((e25 = h26(e15)) | (e25 = h26(e14)) | (e25 = h26(e13)) | (e25 = h26(e12)) | (e25 = h26(e11)) | (e25 = h26(e10))) & ((e24 = h26(e15)) | (e24 = h26(e14)) | (e24 = h26(e13)) | (e24 = h26(e12)) | (e24 = h26(e11)) | (e24 = h26(e10))) & ((e23 = h26(e15)) | (e23 = h26(e14)) | (e23 = h26(e13)) | (e23 = h26(e12)) | (e23 = h26(e11)) | (e23 = h26(e10))) & ((e22 = h26(e15)) | (e22 = h26(e14)) | (e22 = h26(e13)) | (e22 = h26(e12)) | (e22 = h26(e11)) | (e22 = h26(e10))) & ((e21 = h26(e15)) | (e21 = h26(e14)) | (e21 = h26(e13)) | (e21 = h26(e12)) | (e21 = h26(e11)) | (e21 = h26(e10))) & ((e20 = h26(e15)) | (e20 = h26(e14)) | (e20 = h26(e13)) | (e20 = h26(e12)) | (e20 = h26(e11)) | (e20 = h26(e10))) & (h26(op1(e15, e15)) = op2(h26(e15), h26(e15))) & (h26(op1(e15, e14)) = op2(h26(e15), h26(e14))) & (h26(op1(e15, e13)) = op2(h26(e15), h26(e13))) & (h26(op1(e15, e12)) = op2(h26(e15), h26(e12))) & (h26(op1(e15, e11)) = op2(h26(e15), h26(e11))) & (h26(op1(e15, e10)) = op2(h26(e15), h26(e10))) & (h26(op1(e14, e15)) = op2(h26(e14), h26(e15))) & (h26(op1(e14, e14)) = op2(h26(e14), h26(e14))) & (h26(op1(e14, e13)) = op2(h26(e14), h26(e13))) & (h26(op1(e14, e12)) = op2(h26(e14), h26(e12))) & (h26(op1(e14, e11)) = op2(h26(e14), h26(e11))) & (h26(op1(e14, e10)) = op2(h26(e14), h26(e10))) & (h26(op1(e13, e15)) = op2(h26(e13), h26(e15))) & (h26(op1(e13, e14)) = op2(h26(e13), h26(e14))) & (h26(op1(e13, e13)) = op2(h26(e13), h26(e13))) & (h26(op1(e13, e12)) = op2(h26(e13), h26(e12))) & (h26(op1(e13, e11)) = op2(h26(e13), h26(e11))) & (h26(op1(e13, e10)) = op2(h26(e13), h26(e10))) & (h26(op1(e12, e15)) = op2(h26(e12), h26(e15))) & (h26(op1(e12, e14)) = op2(h26(e12), h26(e14))) & (h26(op1(e12, e13)) = op2(h26(e12), h26(e13))) & (h26(op1(e12, e12)) = op2(h26(e12), h26(e12))) & (h26(op1(e12, e11)) = op2(h26(e12), h26(e11))) & (h26(op1(e12, e10)) = op2(h26(e12), h26(e10))) & (h26(op1(e11, e15)) = op2(h26(e11), h26(e15))) & (h26(op1(e11, e14)) = op2(h26(e11), h26(e14))) & (h26(op1(e11, e13)) = op2(h26(e11), h26(e13))) & (h26(op1(e11, e12)) = op2(h26(e11), h26(e12))) & (h26(op1(e11, e11)) = op2(h26(e11), h26(e11))) & (h26(op1(e11, e10)) = op2(h26(e11), h26(e10))) & (h26(op1(e10, e15)) = op2(h26(e10), h26(e15))) & (h26(op1(e10, e14)) = op2(h26(e10), h26(e14))) & (h26(op1(e10, e13)) = op2(h26(e10), h26(e13))) & (h26(op1(e10, e12)) = op2(h26(e10), h26(e12))) & (h26(op1(e10, e11)) = op2(h26(e10), h26(e11))) & (h26(op1(e10, e10)) = op2(h26(e10), h26(e10)))) | (((e25 = h25(e15)) | (e25 = h25(e14)) | (e25 = h25(e13)) | (e25 = h25(e12)) | (e25 = h25(e11)) | (e25 = h25(e10))) & ((e24 = h25(e15)) | (e24 = h25(e14)) | (e24 = h25(e13)) | (e24 = h25(e12)) | (e24 = h25(e11)) | (e24 = h25(e10))) & ((e23 = h25(e15)) | (e23 = h25(e14)) | (e23 = h25(e13)) | (e23 = h25(e12)) | (e23 = h25(e11)) | (e23 = h25(e10))) & ((e22 = h25(e15)) | (e22 = h25(e14)) | (e22 = h25(e13)) | (e22 = h25(e12)) | (e22 = h25(e11)) | (e22 = h25(e10))) & ((e21 = h25(e15)) | (e21 = h25(e14)) | (e21 = h25(e13)) | (e21 = h25(e12)) | (e21 = h25(e11)) | (e21 = h25(e10))) & ((e20 = h25(e15)) | (e20 = h25(e14)) | (e20 = h25(e13)) | (e20 = h25(e12)) | (e20 = h25(e11)) | (e20 = h25(e10))) & (h25(op1(e15, e15)) = op2(h25(e15), h25(e15))) & (h25(op1(e15, e14)) = op2(h25(e15), h25(e14))) & (h25(op1(e15, e13)) = op2(h25(e15), h25(e13))) & (h25(op1(e15, e12)) = op2(h25(e15), h25(e12))) & (h25(op1(e15, e11)) = op2(h25(e15), h25(e11))) & (h25(op1(e15, e10)) = op2(h25(e15), h25(e10))) & (h25(op1(e14, e15)) = op2(h25(e14), h25(e15))) & (h25(op1(e14, e14)) = op2(h25(e14), h25(e14))) & (h25(op1(e14, e13)) = op2(h25(e14), h25(e13))) & (h25(op1(e14, e12)) = op2(h25(e14), h25(e12))) & (h25(op1(e14, e11)) = op2(h25(e14), h25(e11))) & (h25(op1(e14, e10)) = op2(h25(e14), h25(e10))) & (h25(op1(e13, e15)) = op2(h25(e13), h25(e15))) & (h25(op1(e13, e14)) = op2(h25(e13), h25(e14))) & (h25(op1(e13, e13)) = op2(h25(e13), h25(e13))) & (h25(op1(e13, e12)) = op2(h25(e13), h25(e12))) & (h25(op1(e13, e11)) = op2(h25(e13), h25(e11))) & (h25(op1(e13, e10)) = op2(h25(e13), h25(e10))) & (h25(op1(e12, e15)) = op2(h25(e12), h25(e15))) & (h25(op1(e12, e14)) = op2(h25(e12), h25(e14))) & (h25(op1(e12, e13)) = op2(h25(e12), h25(e13))) & (h25(op1(e12, e12)) = op2(h25(e12), h25(e12))) & (h25(op1(e12, e11)) = op2(h25(e12), h25(e11))) & (h25(op1(e12, e10)) = op2(h25(e12), h25(e10))) & (h25(op1(e11, e15)) = op2(h25(e11), h25(e15))) & (h25(op1(e11, e14)) = op2(h25(e11), h25(e14))) & (h25(op1(e11, e13)) = op2(h25(e11), h25(e13))) & (h25(op1(e11, e12)) = op2(h25(e11), h25(e12))) & (h25(op1(e11, e11)) = op2(h25(e11), h25(e11))) & (h25(op1(e11, e10)) = op2(h25(e11), h25(e10))) & (h25(op1(e10, e15)) = op2(h25(e10), h25(e15))) & (h25(op1(e10, e14)) = op2(h25(e10), h25(e14))) & (h25(op1(e10, e13)) = op2(h25(e10), h25(e13))) & (h25(op1(e10, e12)) = op2(h25(e10), h25(e12))) & (h25(op1(e10, e11)) = op2(h25(e10), h25(e11))) & (h25(op1(e10, e10)) = op2(h25(e10), h25(e10)))) | (((e25 = h24(e15)) | (e25 = h24(e14)) | (e25 = h24(e13)) | (e25 = h24(e12)) | (e25 = h24(e11)) | (e25 = h24(e10))) & ((e24 = h24(e15)) | (e24 = h24(e14)) | (e24 = h24(e13)) | (e24 = h24(e12)) | (e24 = h24(e11)) | (e24 = h24(e10))) & ((e23 = h24(e15)) | (e23 = h24(e14)) | (e23 = h24(e13)) | (e23 = h24(e12)) | (e23 = h24(e11)) | (e23 = h24(e10))) & ((e22 = h24(e15)) | (e22 = h24(e14)) | (e22 = h24(e13)) | (e22 = h24(e12)) | (e22 = h24(e11)) | (e22 = h24(e10))) & ((e21 = h24(e15)) | (e21 = h24(e14)) | (e21 = h24(e13)) | (e21 = h24(e12)) | (e21 = h24(e11)) | (e21 = h24(e10))) & ((e20 = h24(e15)) | (e20 = h24(e14)) | (e20 = h24(e13)) | (e20 = h24(e12)) | (e20 = h24(e11)) | (e20 = h24(e10))) & (h24(op1(e15, e15)) = op2(h24(e15), h24(e15))) & (h24(op1(e15, e14)) = op2(h24(e15), h24(e14))) & (h24(op1(e15, e13)) = op2(h24(e15), h24(e13))) & (h24(op1(e15, e12)) = op2(h24(e15), h24(e12))) & (h24(op1(e15, e11)) = op2(h24(e15), h24(e11))) & (h24(op1(e15, e10)) = op2(h24(e15), h24(e10))) & (h24(op1(e14, e15)) = op2(h24(e14), h24(e15))) & (h24(op1(e14, e14)) = op2(h24(e14), h24(e14))) & (h24(op1(e14, e13)) = op2(h24(e14), h24(e13))) & (h24(op1(e14, e12)) = op2(h24(e14), h24(e12))) & (h24(op1(e14, e11)) = op2(h24(e14), h24(e11))) & (h24(op1(e14, e10)) = op2(h24(e14), h24(e10))) & (h24(op1(e13, e15)) = op2(h24(e13), h24(e15))) & (h24(op1(e13, e14)) = op2(h24(e13), h24(e14))) & (h24(op1(e13, e13)) = op2(h24(e13), h24(e13))) & (h24(op1(e13, e12)) = op2(h24(e13), h24(e12))) & (h24(op1(e13, e11)) = op2(h24(e13), h24(e11))) & (h24(op1(e13, e10)) = op2(h24(e13), h24(e10))) & (h24(op1(e12, e15)) = op2(h24(e12), h24(e15))) & (h24(op1(e12, e14)) = op2(h24(e12), h24(e14))) & (h24(op1(e12, e13)) = op2(h24(e12), h24(e13))) & (h24(op1(e12, e12)) = op2(h24(e12), h24(e12))) & (h24(op1(e12, e11)) = op2(h24(e12), h24(e11))) & (h24(op1(e12, e10)) = op2(h24(e12), h24(e10))) & (h24(op1(e11, e15)) = op2(h24(e11), h24(e15))) & (h24(op1(e11, e14)) = op2(h24(e11), h24(e14))) & (h24(op1(e11, e13)) = op2(h24(e11), h24(e13))) & (h24(op1(e11, e12)) = op2(h24(e11), h24(e12))) & (h24(op1(e11, e11)) = op2(h24(e11), h24(e11))) & (h24(op1(e11, e10)) = op2(h24(e11), h24(e10))) & (h24(op1(e10, e15)) = op2(h24(e10), h24(e15))) & (h24(op1(e10, e14)) = op2(h24(e10), h24(e14))) & (h24(op1(e10, e13)) = op2(h24(e10), h24(e13))) & (h24(op1(e10, e12)) = op2(h24(e10), h24(e12))) & (h24(op1(e10, e11)) = op2(h24(e10), h24(e11))) & (h24(op1(e10, e10)) = op2(h24(e10), h24(e10)))) | (((e25 = h23(e15)) | (e25 = h23(e14)) | (e25 = h23(e13)) | (e25 = h23(e12)) | (e25 = h23(e11)) | (e25 = h23(e10))) & ((e24 = h23(e15)) | (e24 = h23(e14)) | (e24 = h23(e13)) | (e24 = h23(e12)) | (e24 = h23(e11)) | (e24 = h23(e10))) & ((e23 = h23(e15)) | (e23 = h23(e14)) | (e23 = h23(e13)) | (e23 = h23(e12)) | (e23 = h23(e11)) | (e23 = h23(e10))) & ((e22 = h23(e15)) | (e22 = h23(e14)) | (e22 = h23(e13)) | (e22 = h23(e12)) | (e22 = h23(e11)) | (e22 = h23(e10))) & ((e21 = h23(e15)) | (e21 = h23(e14)) | (e21 = h23(e13)) | (e21 = h23(e12)) | (e21 = h23(e11)) | (e21 = h23(e10))) & ((e20 = h23(e15)) | (e20 = h23(e14)) | (e20 = h23(e13)) | (e20 = h23(e12)) | (e20 = h23(e11)) | (e20 = h23(e10))) & (h23(op1(e15, e15)) = op2(h23(e15), h23(e15))) & (h23(op1(e15, e14)) = op2(h23(e15), h23(e14))) & (h23(op1(e15, e13)) = op2(h23(e15), h23(e13))) & (h23(op1(e15, e12)) = op2(h23(e15), h23(e12))) & (h23(op1(e15, e11)) = op2(h23(e15), h23(e11))) & (h23(op1(e15, e10)) = op2(h23(e15), h23(e10))) & (h23(op1(e14, e15)) = op2(h23(e14), h23(e15))) & (h23(op1(e14, e14)) = op2(h23(e14), h23(e14))) & (h23(op1(e14, e13)) = op2(h23(e14), h23(e13))) & (h23(op1(e14, e12)) = op2(h23(e14), h23(e12))) & (h23(op1(e14, e11)) = op2(h23(e14), h23(e11))) & (h23(op1(e14, e10)) = op2(h23(e14), h23(e10))) & (h23(op1(e13, e15)) = op2(h23(e13), h23(e15))) & (h23(op1(e13, e14)) = op2(h23(e13), h23(e14))) & (h23(op1(e13, e13)) = op2(h23(e13), h23(e13))) & (h23(op1(e13, e12)) = op2(h23(e13), h23(e12))) & (h23(op1(e13, e11)) = op2(h23(e13), h23(e11))) & (h23(op1(e13, e10)) = op2(h23(e13), h23(e10))) & (h23(op1(e12, e15)) = op2(h23(e12), h23(e15))) & (h23(op1(e12, e14)) = op2(h23(e12), h23(e14))) & (h23(op1(e12, e13)) = op2(h23(e12), h23(e13))) & (h23(op1(e12, e12)) = op2(h23(e12), h23(e12))) & (h23(op1(e12, e11)) = op2(h23(e12), h23(e11))) & (h23(op1(e12, e10)) = op2(h23(e12), h23(e10))) & (h23(op1(e11, e15)) = op2(h23(e11), h23(e15))) & (h23(op1(e11, e14)) = op2(h23(e11), h23(e14))) & (h23(op1(e11, e13)) = op2(h23(e11), h23(e13))) & (h23(op1(e11, e12)) = op2(h23(e11), h23(e12))) & (h23(op1(e11, e11)) = op2(h23(e11), h23(e11))) & (h23(op1(e11, e10)) = op2(h23(e11), h23(e10))) & (h23(op1(e10, e15)) = op2(h23(e10), h23(e15))) & (h23(op1(e10, e14)) = op2(h23(e10), h23(e14))) & (h23(op1(e10, e13)) = op2(h23(e10), h23(e13))) & (h23(op1(e10, e12)) = op2(h23(e10), h23(e12))) & (h23(op1(e10, e11)) = op2(h23(e10), h23(e11))) & (h23(op1(e10, e10)) = op2(h23(e10), h23(e10)))) | (((e25 = h22(e15)) | (e25 = h22(e14)) | (e25 = h22(e13)) | (e25 = h22(e12)) | (e25 = h22(e11)) | (e25 = h22(e10))) & ((e24 = h22(e15)) | (e24 = h22(e14)) | (e24 = h22(e13)) | (e24 = h22(e12)) | (e24 = h22(e11)) | (e24 = h22(e10))) & ((e23 = h22(e15)) | (e23 = h22(e14)) | (e23 = h22(e13)) | (e23 = h22(e12)) | (e23 = h22(e11)) | (e23 = h22(e10))) & ((e22 = h22(e15)) | (e22 = h22(e14)) | (e22 = h22(e13)) | (e22 = h22(e12)) | (e22 = h22(e11)) | (e22 = h22(e10))) & ((e21 = h22(e15)) | (e21 = h22(e14)) | (e21 = h22(e13)) | (e21 = h22(e12)) | (e21 = h22(e11)) | (e21 = h22(e10))) & ((e20 = h22(e15)) | (e20 = h22(e14)) | (e20 = h22(e13)) | (e20 = h22(e12)) | (e20 = h22(e11)) | (e20 = h22(e10))) & (h22(op1(e15, e15)) = op2(h22(e15), h22(e15))) & (h22(op1(e15, e14)) = op2(h22(e15), h22(e14))) & (h22(op1(e15, e13)) = op2(h22(e15), h22(e13))) & (h22(op1(e15, e12)) = op2(h22(e15), h22(e12))) & (h22(op1(e15, e11)) = op2(h22(e15), h22(e11))) & (h22(op1(e15, e10)) = op2(h22(e15), h22(e10))) & (h22(op1(e14, e15)) = op2(h22(e14), h22(e15))) & (h22(op1(e14, e14)) = op2(h22(e14), h22(e14))) & (h22(op1(e14, e13)) = op2(h22(e14), h22(e13))) & (h22(op1(e14, e12)) = op2(h22(e14), h22(e12))) & (h22(op1(e14, e11)) = op2(h22(e14), h22(e11))) & (h22(op1(e14, e10)) = op2(h22(e14), h22(e10))) & (h22(op1(e13, e15)) = op2(h22(e13), h22(e15))) & (h22(op1(e13, e14)) = op2(h22(e13), h22(e14))) & (h22(op1(e13, e13)) = op2(h22(e13), h22(e13))) & (h22(op1(e13, e12)) = op2(h22(e13), h22(e12))) & (h22(op1(e13, e11)) = op2(h22(e13), h22(e11))) & (h22(op1(e13, e10)) = op2(h22(e13), h22(e10))) & (h22(op1(e12, e15)) = op2(h22(e12), h22(e15))) & (h22(op1(e12, e14)) = op2(h22(e12), h22(e14))) & (h22(op1(e12, e13)) = op2(h22(e12), h22(e13))) & (h22(op1(e12, e12)) = op2(h22(e12), h22(e12))) & (h22(op1(e12, e11)) = op2(h22(e12), h22(e11))) & (h22(op1(e12, e10)) = op2(h22(e12), h22(e10))) & (h22(op1(e11, e15)) = op2(h22(e11), h22(e15))) & (h22(op1(e11, e14)) = op2(h22(e11), h22(e14))) & (h22(op1(e11, e13)) = op2(h22(e11), h22(e13))) & (h22(op1(e11, e12)) = op2(h22(e11), h22(e12))) & (h22(op1(e11, e11)) = op2(h22(e11), h22(e11))) & (h22(op1(e11, e10)) = op2(h22(e11), h22(e10))) & (h22(op1(e10, e15)) = op2(h22(e10), h22(e15))) & (h22(op1(e10, e14)) = op2(h22(e10), h22(e14))) & (h22(op1(e10, e13)) = op2(h22(e10), h22(e13))) & (h22(op1(e10, e12)) = op2(h22(e10), h22(e12))) & (h22(op1(e10, e11)) = op2(h22(e10), h22(e11))) & (h22(op1(e10, e10)) = op2(h22(e10), h22(e10)))) | (((e25 = h21(e15)) | (e25 = h21(e14)) | (e25 = h21(e13)) | (e25 = h21(e12)) | (e25 = h21(e11)) | (e25 = h21(e10))) & ((e24 = h21(e15)) | (e24 = h21(e14)) | (e24 = h21(e13)) | (e24 = h21(e12)) | (e24 = h21(e11)) | (e24 = h21(e10))) & ((e23 = h21(e15)) | (e23 = h21(e14)) | (e23 = h21(e13)) | (e23 = h21(e12)) | (e23 = h21(e11)) | (e23 = h21(e10))) & ((e22 = h21(e15)) | (e22 = h21(e14)) | (e22 = h21(e13)) | (e22 = h21(e12)) | (e22 = h21(e11)) | (e22 = h21(e10))) & ((e21 = h21(e15)) | (e21 = h21(e14)) | (e21 = h21(e13)) | (e21 = h21(e12)) | (e21 = h21(e11)) | (e21 = h21(e10))) & ((e20 = h21(e15)) | (e20 = h21(e14)) | (e20 = h21(e13)) | (e20 = h21(e12)) | (e20 = h21(e11)) | (e20 = h21(e10))) & (h21(op1(e15, e15)) = op2(h21(e15), h21(e15))) & (h21(op1(e15, e14)) = op2(h21(e15), h21(e14))) & (h21(op1(e15, e13)) = op2(h21(e15), h21(e13))) & (h21(op1(e15, e12)) = op2(h21(e15), h21(e12))) & (h21(op1(e15, e11)) = op2(h21(e15), h21(e11))) & (h21(op1(e15, e10)) = op2(h21(e15), h21(e10))) & (h21(op1(e14, e15)) = op2(h21(e14), h21(e15))) & (h21(op1(e14, e14)) = op2(h21(e14), h21(e14))) & (h21(op1(e14, e13)) = op2(h21(e14), h21(e13))) & (h21(op1(e14, e12)) = op2(h21(e14), h21(e12))) & (h21(op1(e14, e11)) = op2(h21(e14), h21(e11))) & (h21(op1(e14, e10)) = op2(h21(e14), h21(e10))) & (h21(op1(e13, e15)) = op2(h21(e13), h21(e15))) & (h21(op1(e13, e14)) = op2(h21(e13), h21(e14))) & (h21(op1(e13, e13)) = op2(h21(e13), h21(e13))) & (h21(op1(e13, e12)) = op2(h21(e13), h21(e12))) & (h21(op1(e13, e11)) = op2(h21(e13), h21(e11))) & (h21(op1(e13, e10)) = op2(h21(e13), h21(e10))) & (h21(op1(e12, e15)) = op2(h21(e12), h21(e15))) & (h21(op1(e12, e14)) = op2(h21(e12), h21(e14))) & (h21(op1(e12, e13)) = op2(h21(e12), h21(e13))) & (h21(op1(e12, e12)) = op2(h21(e12), h21(e12))) & (h21(op1(e12, e11)) = op2(h21(e12), h21(e11))) & (h21(op1(e12, e10)) = op2(h21(e12), h21(e10))) & (h21(op1(e11, e15)) = op2(h21(e11), h21(e15))) & (h21(op1(e11, e14)) = op2(h21(e11), h21(e14))) & (h21(op1(e11, e13)) = op2(h21(e11), h21(e13))) & (h21(op1(e11, e12)) = op2(h21(e11), h21(e12))) & (h21(op1(e11, e11)) = op2(h21(e11), h21(e11))) & (h21(op1(e11, e10)) = op2(h21(e11), h21(e10))) & (h21(op1(e10, e15)) = op2(h21(e10), h21(e15))) & (h21(op1(e10, e14)) = op2(h21(e10), h21(e14))) & (h21(op1(e10, e13)) = op2(h21(e10), h21(e13))) & (h21(op1(e10, e12)) = op2(h21(e10), h21(e12))) & (h21(op1(e10, e11)) = op2(h21(e10), h21(e11))) & (h21(op1(e10, e10)) = op2(h21(e10), h21(e10)))) | (((e25 = h20(e15)) | (e25 = h20(e14)) | (e25 = h20(e13)) | (e25 = h20(e12)) | (e25 = h20(e11)) | (e25 = h20(e10))) & ((e24 = h20(e15)) | (e24 = h20(e14)) | (e24 = h20(e13)) | (e24 = h20(e12)) | (e24 = h20(e11)) | (e24 = h20(e10))) & ((e23 = h20(e15)) | (e23 = h20(e14)) | (e23 = h20(e13)) | (e23 = h20(e12)) | (e23 = h20(e11)) | (e23 = h20(e10))) & ((e22 = h20(e15)) | (e22 = h20(e14)) | (e22 = h20(e13)) | (e22 = h20(e12)) | (e22 = h20(e11)) | (e22 = h20(e10))) & ((e21 = h20(e15)) | (e21 = h20(e14)) | (e21 = h20(e13)) | (e21 = h20(e12)) | (e21 = h20(e11)) | (e21 = h20(e10))) & ((e20 = h20(e15)) | (e20 = h20(e14)) | (e20 = h20(e13)) | (e20 = h20(e12)) | (e20 = h20(e11)) | (e20 = h20(e10))) & (h20(op1(e15, e15)) = op2(h20(e15), h20(e15))) & (h20(op1(e15, e14)) = op2(h20(e15), h20(e14))) & (h20(op1(e15, e13)) = op2(h20(e15), h20(e13))) & (h20(op1(e15, e12)) = op2(h20(e15), h20(e12))) & (h20(op1(e15, e11)) = op2(h20(e15), h20(e11))) & (h20(op1(e15, e10)) = op2(h20(e15), h20(e10))) & (h20(op1(e14, e15)) = op2(h20(e14), h20(e15))) & (h20(op1(e14, e14)) = op2(h20(e14), h20(e14))) & (h20(op1(e14, e13)) = op2(h20(e14), h20(e13))) & (h20(op1(e14, e12)) = op2(h20(e14), h20(e12))) & (h20(op1(e14, e11)) = op2(h20(e14), h20(e11))) & (h20(op1(e14, e10)) = op2(h20(e14), h20(e10))) & (h20(op1(e13, e15)) = op2(h20(e13), h20(e15))) & (h20(op1(e13, e14)) = op2(h20(e13), h20(e14))) & (h20(op1(e13, e13)) = op2(h20(e13), h20(e13))) & (h20(op1(e13, e12)) = op2(h20(e13), h20(e12))) & (h20(op1(e13, e11)) = op2(h20(e13), h20(e11))) & (h20(op1(e13, e10)) = op2(h20(e13), h20(e10))) & (h20(op1(e12, e15)) = op2(h20(e12), h20(e15))) & (h20(op1(e12, e14)) = op2(h20(e12), h20(e14))) & (h20(op1(e12, e13)) = op2(h20(e12), h20(e13))) & (h20(op1(e12, e12)) = op2(h20(e12), h20(e12))) & (h20(op1(e12, e11)) = op2(h20(e12), h20(e11))) & (h20(op1(e12, e10)) = op2(h20(e12), h20(e10))) & (h20(op1(e11, e15)) = op2(h20(e11), h20(e15))) & (h20(op1(e11, e14)) = op2(h20(e11), h20(e14))) & (h20(op1(e11, e13)) = op2(h20(e11), h20(e13))) & (h20(op1(e11, e12)) = op2(h20(e11), h20(e12))) & (h20(op1(e11, e11)) = op2(h20(e11), h20(e11))) & (h20(op1(e11, e10)) = op2(h20(e11), h20(e10))) & (h20(op1(e10, e15)) = op2(h20(e10), h20(e15))) & (h20(op1(e10, e14)) = op2(h20(e10), h20(e14))) & (h20(op1(e10, e13)) = op2(h20(e10), h20(e13))) & (h20(op1(e10, e12)) = op2(h20(e10), h20(e12))) & (h20(op1(e10, e11)) = op2(h20(e10), h20(e11))) & (h20(op1(e10, e10)) = op2(h20(e10), h20(e10)))) | (((e25 = h19(e15)) | (e25 = h19(e14)) | (e25 = h19(e13)) | (e25 = h19(e12)) | (e25 = h19(e11)) | (e25 = h19(e10))) & ((e24 = h19(e15)) | (e24 = h19(e14)) | (e24 = h19(e13)) | (e24 = h19(e12)) | (e24 = h19(e11)) | (e24 = h19(e10))) & ((e23 = h19(e15)) | (e23 = h19(e14)) | (e23 = h19(e13)) | (e23 = h19(e12)) | (e23 = h19(e11)) | (e23 = h19(e10))) & ((e22 = h19(e15)) | (e22 = h19(e14)) | (e22 = h19(e13)) | (e22 = h19(e12)) | (e22 = h19(e11)) | (e22 = h19(e10))) & ((e21 = h19(e15)) | (e21 = h19(e14)) | (e21 = h19(e13)) | (e21 = h19(e12)) | (e21 = h19(e11)) | (e21 = h19(e10))) & ((e20 = h19(e15)) | (e20 = h19(e14)) | (e20 = h19(e13)) | (e20 = h19(e12)) | (e20 = h19(e11)) | (e20 = h19(e10))) & (h19(op1(e15, e15)) = op2(h19(e15), h19(e15))) & (h19(op1(e15, e14)) = op2(h19(e15), h19(e14))) & (h19(op1(e15, e13)) = op2(h19(e15), h19(e13))) & (h19(op1(e15, e12)) = op2(h19(e15), h19(e12))) & (h19(op1(e15, e11)) = op2(h19(e15), h19(e11))) & (h19(op1(e15, e10)) = op2(h19(e15), h19(e10))) & (h19(op1(e14, e15)) = op2(h19(e14), h19(e15))) & (h19(op1(e14, e14)) = op2(h19(e14), h19(e14))) & (h19(op1(e14, e13)) = op2(h19(e14), h19(e13))) & (h19(op1(e14, e12)) = op2(h19(e14), h19(e12))) & (h19(op1(e14, e11)) = op2(h19(e14), h19(e11))) & (h19(op1(e14, e10)) = op2(h19(e14), h19(e10))) & (h19(op1(e13, e15)) = op2(h19(e13), h19(e15))) & (h19(op1(e13, e14)) = op2(h19(e13), h19(e14))) & (h19(op1(e13, e13)) = op2(h19(e13), h19(e13))) & (h19(op1(e13, e12)) = op2(h19(e13), h19(e12))) & (h19(op1(e13, e11)) = op2(h19(e13), h19(e11))) & (h19(op1(e13, e10)) = op2(h19(e13), h19(e10))) & (h19(op1(e12, e15)) = op2(h19(e12), h19(e15))) & (h19(op1(e12, e14)) = op2(h19(e12), h19(e14))) & (h19(op1(e12, e13)) = op2(h19(e12), h19(e13))) & (h19(op1(e12, e12)) = op2(h19(e12), h19(e12))) & (h19(op1(e12, e11)) = op2(h19(e12), h19(e11))) & (h19(op1(e12, e10)) = op2(h19(e12), h19(e10))) & (h19(op1(e11, e15)) = op2(h19(e11), h19(e15))) & (h19(op1(e11, e14)) = op2(h19(e11), h19(e14))) & (h19(op1(e11, e13)) = op2(h19(e11), h19(e13))) & (h19(op1(e11, e12)) = op2(h19(e11), h19(e12))) & (h19(op1(e11, e11)) = op2(h19(e11), h19(e11))) & (h19(op1(e11, e10)) = op2(h19(e11), h19(e10))) & (h19(op1(e10, e15)) = op2(h19(e10), h19(e15))) & (h19(op1(e10, e14)) = op2(h19(e10), h19(e14))) & (h19(op1(e10, e13)) = op2(h19(e10), h19(e13))) & (h19(op1(e10, e12)) = op2(h19(e10), h19(e12))) & (h19(op1(e10, e11)) = op2(h19(e10), h19(e11))) & (h19(op1(e10, e10)) = op2(h19(e10), h19(e10)))) | (((e25 = h18(e15)) | (e25 = h18(e14)) | (e25 = h18(e13)) | (e25 = h18(e12)) | (e25 = h18(e11)) | (e25 = h18(e10))) & ((e24 = h18(e15)) | (e24 = h18(e14)) | (e24 = h18(e13)) | (e24 = h18(e12)) | (e24 = h18(e11)) | (e24 = h18(e10))) & ((e23 = h18(e15)) | (e23 = h18(e14)) | (e23 = h18(e13)) | (e23 = h18(e12)) | (e23 = h18(e11)) | (e23 = h18(e10))) & ((e22 = h18(e15)) | (e22 = h18(e14)) | (e22 = h18(e13)) | (e22 = h18(e12)) | (e22 = h18(e11)) | (e22 = h18(e10))) & ((e21 = h18(e15)) | (e21 = h18(e14)) | (e21 = h18(e13)) | (e21 = h18(e12)) | (e21 = h18(e11)) | (e21 = h18(e10))) & ((e20 = h18(e15)) | (e20 = h18(e14)) | (e20 = h18(e13)) | (e20 = h18(e12)) | (e20 = h18(e11)) | (e20 = h18(e10))) & (h18(op1(e15, e15)) = op2(h18(e15), h18(e15))) & (h18(op1(e15, e14)) = op2(h18(e15), h18(e14))) & (h18(op1(e15, e13)) = op2(h18(e15), h18(e13))) & (h18(op1(e15, e12)) = op2(h18(e15), h18(e12))) & (h18(op1(e15, e11)) = op2(h18(e15), h18(e11))) & (h18(op1(e15, e10)) = op2(h18(e15), h18(e10))) & (h18(op1(e14, e15)) = op2(h18(e14), h18(e15))) & (h18(op1(e14, e14)) = op2(h18(e14), h18(e14))) & (h18(op1(e14, e13)) = op2(h18(e14), h18(e13))) & (h18(op1(e14, e12)) = op2(h18(e14), h18(e12))) & (h18(op1(e14, e11)) = op2(h18(e14), h18(e11))) & (h18(op1(e14, e10)) = op2(h18(e14), h18(e10))) & (h18(op1(e13, e15)) = op2(h18(e13), h18(e15))) & (h18(op1(e13, e14)) = op2(h18(e13), h18(e14))) & (h18(op1(e13, e13)) = op2(h18(e13), h18(e13))) & (h18(op1(e13, e12)) = op2(h18(e13), h18(e12))) & (h18(op1(e13, e11)) = op2(h18(e13), h18(e11))) & (h18(op1(e13, e10)) = op2(h18(e13), h18(e10))) & (h18(op1(e12, e15)) = op2(h18(e12), h18(e15))) & (h18(op1(e12, e14)) = op2(h18(e12), h18(e14))) & (h18(op1(e12, e13)) = op2(h18(e12), h18(e13))) & (h18(op1(e12, e12)) = op2(h18(e12), h18(e12))) & (h18(op1(e12, e11)) = op2(h18(e12), h18(e11))) & (h18(op1(e12, e10)) = op2(h18(e12), h18(e10))) & (h18(op1(e11, e15)) = op2(h18(e11), h18(e15))) & (h18(op1(e11, e14)) = op2(h18(e11), h18(e14))) & (h18(op1(e11, e13)) = op2(h18(e11), h18(e13))) & (h18(op1(e11, e12)) = op2(h18(e11), h18(e12))) & (h18(op1(e11, e11)) = op2(h18(e11), h18(e11))) & (h18(op1(e11, e10)) = op2(h18(e11), h18(e10))) & (h18(op1(e10, e15)) = op2(h18(e10), h18(e15))) & (h18(op1(e10, e14)) = op2(h18(e10), h18(e14))) & (h18(op1(e10, e13)) = op2(h18(e10), h18(e13))) & (h18(op1(e10, e12)) = op2(h18(e10), h18(e12))) & (h18(op1(e10, e11)) = op2(h18(e10), h18(e11))) & (h18(op1(e10, e10)) = op2(h18(e10), h18(e10)))) | (((e25 = h17(e15)) | (e25 = h17(e14)) | (e25 = h17(e13)) | (e25 = h17(e12)) | (e25 = h17(e11)) | (e25 = h17(e10))) & ((e24 = h17(e15)) | (e24 = h17(e14)) | (e24 = h17(e13)) | (e24 = h17(e12)) | (e24 = h17(e11)) | (e24 = h17(e10))) & ((e23 = h17(e15)) | (e23 = h17(e14)) | (e23 = h17(e13)) | (e23 = h17(e12)) | (e23 = h17(e11)) | (e23 = h17(e10))) & ((e22 = h17(e15)) | (e22 = h17(e14)) | (e22 = h17(e13)) | (e22 = h17(e12)) | (e22 = h17(e11)) | (e22 = h17(e10))) & ((e21 = h17(e15)) | (e21 = h17(e14)) | (e21 = h17(e13)) | (e21 = h17(e12)) | (e21 = h17(e11)) | (e21 = h17(e10))) & ((e20 = h17(e15)) | (e20 = h17(e14)) | (e20 = h17(e13)) | (e20 = h17(e12)) | (e20 = h17(e11)) | (e20 = h17(e10))) & (h17(op1(e15, e15)) = op2(h17(e15), h17(e15))) & (h17(op1(e15, e14)) = op2(h17(e15), h17(e14))) & (h17(op1(e15, e13)) = op2(h17(e15), h17(e13))) & (h17(op1(e15, e12)) = op2(h17(e15), h17(e12))) & (h17(op1(e15, e11)) = op2(h17(e15), h17(e11))) & (h17(op1(e15, e10)) = op2(h17(e15), h17(e10))) & (h17(op1(e14, e15)) = op2(h17(e14), h17(e15))) & (h17(op1(e14, e14)) = op2(h17(e14), h17(e14))) & (h17(op1(e14, e13)) = op2(h17(e14), h17(e13))) & (h17(op1(e14, e12)) = op2(h17(e14), h17(e12))) & (h17(op1(e14, e11)) = op2(h17(e14), h17(e11))) & (h17(op1(e14, e10)) = op2(h17(e14), h17(e10))) & (h17(op1(e13, e15)) = op2(h17(e13), h17(e15))) & (h17(op1(e13, e14)) = op2(h17(e13), h17(e14))) & (h17(op1(e13, e13)) = op2(h17(e13), h17(e13))) & (h17(op1(e13, e12)) = op2(h17(e13), h17(e12))) & (h17(op1(e13, e11)) = op2(h17(e13), h17(e11))) & (h17(op1(e13, e10)) = op2(h17(e13), h17(e10))) & (h17(op1(e12, e15)) = op2(h17(e12), h17(e15))) & (h17(op1(e12, e14)) = op2(h17(e12), h17(e14))) & (h17(op1(e12, e13)) = op2(h17(e12), h17(e13))) & (h17(op1(e12, e12)) = op2(h17(e12), h17(e12))) & (h17(op1(e12, e11)) = op2(h17(e12), h17(e11))) & (h17(op1(e12, e10)) = op2(h17(e12), h17(e10))) & (h17(op1(e11, e15)) = op2(h17(e11), h17(e15))) & (h17(op1(e11, e14)) = op2(h17(e11), h17(e14))) & (h17(op1(e11, e13)) = op2(h17(e11), h17(e13))) & (h17(op1(e11, e12)) = op2(h17(e11), h17(e12))) & (h17(op1(e11, e11)) = op2(h17(e11), h17(e11))) & (h17(op1(e11, e10)) = op2(h17(e11), h17(e10))) & (h17(op1(e10, e15)) = op2(h17(e10), h17(e15))) & (h17(op1(e10, e14)) = op2(h17(e10), h17(e14))) & (h17(op1(e10, e13)) = op2(h17(e10), h17(e13))) & (h17(op1(e10, e12)) = op2(h17(e10), h17(e12))) & (h17(op1(e10, e11)) = op2(h17(e10), h17(e11))) & (h17(op1(e10, e10)) = op2(h17(e10), h17(e10)))) | (((e25 = h16(e15)) | (e25 = h16(e14)) | (e25 = h16(e13)) | (e25 = h16(e12)) | (e25 = h16(e11)) | (e25 = h16(e10))) & ((e24 = h16(e15)) | (e24 = h16(e14)) | (e24 = h16(e13)) | (e24 = h16(e12)) | (e24 = h16(e11)) | (e24 = h16(e10))) & ((e23 = h16(e15)) | (e23 = h16(e14)) | (e23 = h16(e13)) | (e23 = h16(e12)) | (e23 = h16(e11)) | (e23 = h16(e10))) & ((e22 = h16(e15)) | (e22 = h16(e14)) | (e22 = h16(e13)) | (e22 = h16(e12)) | (e22 = h16(e11)) | (e22 = h16(e10))) & ((e21 = h16(e15)) | (e21 = h16(e14)) | (e21 = h16(e13)) | (e21 = h16(e12)) | (e21 = h16(e11)) | (e21 = h16(e10))) & ((e20 = h16(e15)) | (e20 = h16(e14)) | (e20 = h16(e13)) | (e20 = h16(e12)) | (e20 = h16(e11)) | (e20 = h16(e10))) & (h16(op1(e15, e15)) = op2(h16(e15), h16(e15))) & (h16(op1(e15, e14)) = op2(h16(e15), h16(e14))) & (h16(op1(e15, e13)) = op2(h16(e15), h16(e13))) & (h16(op1(e15, e12)) = op2(h16(e15), h16(e12))) & (h16(op1(e15, e11)) = op2(h16(e15), h16(e11))) & (h16(op1(e15, e10)) = op2(h16(e15), h16(e10))) & (h16(op1(e14, e15)) = op2(h16(e14), h16(e15))) & (h16(op1(e14, e14)) = op2(h16(e14), h16(e14))) & (h16(op1(e14, e13)) = op2(h16(e14), h16(e13))) & (h16(op1(e14, e12)) = op2(h16(e14), h16(e12))) & (h16(op1(e14, e11)) = op2(h16(e14), h16(e11))) & (h16(op1(e14, e10)) = op2(h16(e14), h16(e10))) & (h16(op1(e13, e15)) = op2(h16(e13), h16(e15))) & (h16(op1(e13, e14)) = op2(h16(e13), h16(e14))) & (h16(op1(e13, e13)) = op2(h16(e13), h16(e13))) & (h16(op1(e13, e12)) = op2(h16(e13), h16(e12))) & (h16(op1(e13, e11)) = op2(h16(e13), h16(e11))) & (h16(op1(e13, e10)) = op2(h16(e13), h16(e10))) & (h16(op1(e12, e15)) = op2(h16(e12), h16(e15))) & (h16(op1(e12, e14)) = op2(h16(e12), h16(e14))) & (h16(op1(e12, e13)) = op2(h16(e12), h16(e13))) & (h16(op1(e12, e12)) = op2(h16(e12), h16(e12))) & (h16(op1(e12, e11)) = op2(h16(e12), h16(e11))) & (h16(op1(e12, e10)) = op2(h16(e12), h16(e10))) & (h16(op1(e11, e15)) = op2(h16(e11), h16(e15))) & (h16(op1(e11, e14)) = op2(h16(e11), h16(e14))) & (h16(op1(e11, e13)) = op2(h16(e11), h16(e13))) & (h16(op1(e11, e12)) = op2(h16(e11), h16(e12))) & (h16(op1(e11, e11)) = op2(h16(e11), h16(e11))) & (h16(op1(e11, e10)) = op2(h16(e11), h16(e10))) & (h16(op1(e10, e15)) = op2(h16(e10), h16(e15))) & (h16(op1(e10, e14)) = op2(h16(e10), h16(e14))) & (h16(op1(e10, e13)) = op2(h16(e10), h16(e13))) & (h16(op1(e10, e12)) = op2(h16(e10), h16(e12))) & (h16(op1(e10, e11)) = op2(h16(e10), h16(e11))) & (h16(op1(e10, e10)) = op2(h16(e10), h16(e10)))) | (((e25 = h15(e15)) | (e25 = h15(e14)) | (e25 = h15(e13)) | (e25 = h15(e12)) | (e25 = h15(e11)) | (e25 = h15(e10))) & ((e24 = h15(e15)) | (e24 = h15(e14)) | (e24 = h15(e13)) | (e24 = h15(e12)) | (e24 = h15(e11)) | (e24 = h15(e10))) & ((e23 = h15(e15)) | (e23 = h15(e14)) | (e23 = h15(e13)) | (e23 = h15(e12)) | (e23 = h15(e11)) | (e23 = h15(e10))) & ((e22 = h15(e15)) | (e22 = h15(e14)) | (e22 = h15(e13)) | (e22 = h15(e12)) | (e22 = h15(e11)) | (e22 = h15(e10))) & ((e21 = h15(e15)) | (e21 = h15(e14)) | (e21 = h15(e13)) | (e21 = h15(e12)) | (e21 = h15(e11)) | (e21 = h15(e10))) & ((e20 = h15(e15)) | (e20 = h15(e14)) | (e20 = h15(e13)) | (e20 = h15(e12)) | (e20 = h15(e11)) | (e20 = h15(e10))) & (h15(op1(e15, e15)) = op2(h15(e15), h15(e15))) & (h15(op1(e15, e14)) = op2(h15(e15), h15(e14))) & (h15(op1(e15, e13)) = op2(h15(e15), h15(e13))) & (h15(op1(e15, e12)) = op2(h15(e15), h15(e12))) & (h15(op1(e15, e11)) = op2(h15(e15), h15(e11))) & (h15(op1(e15, e10)) = op2(h15(e15), h15(e10))) & (h15(op1(e14, e15)) = op2(h15(e14), h15(e15))) & (h15(op1(e14, e14)) = op2(h15(e14), h15(e14))) & (h15(op1(e14, e13)) = op2(h15(e14), h15(e13))) & (h15(op1(e14, e12)) = op2(h15(e14), h15(e12))) & (h15(op1(e14, e11)) = op2(h15(e14), h15(e11))) & (h15(op1(e14, e10)) = op2(h15(e14), h15(e10))) & (h15(op1(e13, e15)) = op2(h15(e13), h15(e15))) & (h15(op1(e13, e14)) = op2(h15(e13), h15(e14))) & (h15(op1(e13, e13)) = op2(h15(e13), h15(e13))) & (h15(op1(e13, e12)) = op2(h15(e13), h15(e12))) & (h15(op1(e13, e11)) = op2(h15(e13), h15(e11))) & (h15(op1(e13, e10)) = op2(h15(e13), h15(e10))) & (h15(op1(e12, e15)) = op2(h15(e12), h15(e15))) & (h15(op1(e12, e14)) = op2(h15(e12), h15(e14))) & (h15(op1(e12, e13)) = op2(h15(e12), h15(e13))) & (h15(op1(e12, e12)) = op2(h15(e12), h15(e12))) & (h15(op1(e12, e11)) = op2(h15(e12), h15(e11))) & (h15(op1(e12, e10)) = op2(h15(e12), h15(e10))) & (h15(op1(e11, e15)) = op2(h15(e11), h15(e15))) & (h15(op1(e11, e14)) = op2(h15(e11), h15(e14))) & (h15(op1(e11, e13)) = op2(h15(e11), h15(e13))) & (h15(op1(e11, e12)) = op2(h15(e11), h15(e12))) & (h15(op1(e11, e11)) = op2(h15(e11), h15(e11))) & (h15(op1(e11, e10)) = op2(h15(e11), h15(e10))) & (h15(op1(e10, e15)) = op2(h15(e10), h15(e15))) & (h15(op1(e10, e14)) = op2(h15(e10), h15(e14))) & (h15(op1(e10, e13)) = op2(h15(e10), h15(e13))) & (h15(op1(e10, e12)) = op2(h15(e10), h15(e12))) & (h15(op1(e10, e11)) = op2(h15(e10), h15(e11))) & (h15(op1(e10, e10)) = op2(h15(e10), h15(e10)))) | (((e25 = h14(e15)) | (e25 = h14(e14)) | (e25 = h14(e13)) | (e25 = h14(e12)) | (e25 = h14(e11)) | (e25 = h14(e10))) & ((e24 = h14(e15)) | (e24 = h14(e14)) | (e24 = h14(e13)) | (e24 = h14(e12)) | (e24 = h14(e11)) | (e24 = h14(e10))) & ((e23 = h14(e15)) | (e23 = h14(e14)) | (e23 = h14(e13)) | (e23 = h14(e12)) | (e23 = h14(e11)) | (e23 = h14(e10))) & ((e22 = h14(e15)) | (e22 = h14(e14)) | (e22 = h14(e13)) | (e22 = h14(e12)) | (e22 = h14(e11)) | (e22 = h14(e10))) & ((e21 = h14(e15)) | (e21 = h14(e14)) | (e21 = h14(e13)) | (e21 = h14(e12)) | (e21 = h14(e11)) | (e21 = h14(e10))) & ((e20 = h14(e15)) | (e20 = h14(e14)) | (e20 = h14(e13)) | (e20 = h14(e12)) | (e20 = h14(e11)) | (e20 = h14(e10))) & (h14(op1(e15, e15)) = op2(h14(e15), h14(e15))) & (h14(op1(e15, e14)) = op2(h14(e15), h14(e14))) & (h14(op1(e15, e13)) = op2(h14(e15), h14(e13))) & (h14(op1(e15, e12)) = op2(h14(e15), h14(e12))) & (h14(op1(e15, e11)) = op2(h14(e15), h14(e11))) & (h14(op1(e15, e10)) = op2(h14(e15), h14(e10))) & (h14(op1(e14, e15)) = op2(h14(e14), h14(e15))) & (h14(op1(e14, e14)) = op2(h14(e14), h14(e14))) & (h14(op1(e14, e13)) = op2(h14(e14), h14(e13))) & (h14(op1(e14, e12)) = op2(h14(e14), h14(e12))) & (h14(op1(e14, e11)) = op2(h14(e14), h14(e11))) & (h14(op1(e14, e10)) = op2(h14(e14), h14(e10))) & (h14(op1(e13, e15)) = op2(h14(e13), h14(e15))) & (h14(op1(e13, e14)) = op2(h14(e13), h14(e14))) & (h14(op1(e13, e13)) = op2(h14(e13), h14(e13))) & (h14(op1(e13, e12)) = op2(h14(e13), h14(e12))) & (h14(op1(e13, e11)) = op2(h14(e13), h14(e11))) & (h14(op1(e13, e10)) = op2(h14(e13), h14(e10))) & (h14(op1(e12, e15)) = op2(h14(e12), h14(e15))) & (h14(op1(e12, e14)) = op2(h14(e12), h14(e14))) & (h14(op1(e12, e13)) = op2(h14(e12), h14(e13))) & (h14(op1(e12, e12)) = op2(h14(e12), h14(e12))) & (h14(op1(e12, e11)) = op2(h14(e12), h14(e11))) & (h14(op1(e12, e10)) = op2(h14(e12), h14(e10))) & (h14(op1(e11, e15)) = op2(h14(e11), h14(e15))) & (h14(op1(e11, e14)) = op2(h14(e11), h14(e14))) & (h14(op1(e11, e13)) = op2(h14(e11), h14(e13))) & (h14(op1(e11, e12)) = op2(h14(e11), h14(e12))) & (h14(op1(e11, e11)) = op2(h14(e11), h14(e11))) & (h14(op1(e11, e10)) = op2(h14(e11), h14(e10))) & (h14(op1(e10, e15)) = op2(h14(e10), h14(e15))) & (h14(op1(e10, e14)) = op2(h14(e10), h14(e14))) & (h14(op1(e10, e13)) = op2(h14(e10), h14(e13))) & (h14(op1(e10, e12)) = op2(h14(e10), h14(e12))) & (h14(op1(e10, e11)) = op2(h14(e10), h14(e11))) & (h14(op1(e10, e10)) = op2(h14(e10), h14(e10)))) | (((e25 = h13(e15)) | (e25 = h13(e14)) | (e25 = h13(e13)) | (e25 = h13(e12)) | (e25 = h13(e11)) | (e25 = h13(e10))) & ((e24 = h13(e15)) | (e24 = h13(e14)) | (e24 = h13(e13)) | (e24 = h13(e12)) | (e24 = h13(e11)) | (e24 = h13(e10))) & ((e23 = h13(e15)) | (e23 = h13(e14)) | (e23 = h13(e13)) | (e23 = h13(e12)) | (e23 = h13(e11)) | (e23 = h13(e10))) & ((e22 = h13(e15)) | (e22 = h13(e14)) | (e22 = h13(e13)) | (e22 = h13(e12)) | (e22 = h13(e11)) | (e22 = h13(e10))) & ((e21 = h13(e15)) | (e21 = h13(e14)) | (e21 = h13(e13)) | (e21 = h13(e12)) | (e21 = h13(e11)) | (e21 = h13(e10))) & ((e20 = h13(e15)) | (e20 = h13(e14)) | (e20 = h13(e13)) | (e20 = h13(e12)) | (e20 = h13(e11)) | (e20 = h13(e10))) & (h13(op1(e15, e15)) = op2(h13(e15), h13(e15))) & (h13(op1(e15, e14)) = op2(h13(e15), h13(e14))) & (h13(op1(e15, e13)) = op2(h13(e15), h13(e13))) & (h13(op1(e15, e12)) = op2(h13(e15), h13(e12))) & (h13(op1(e15, e11)) = op2(h13(e15), h13(e11))) & (h13(op1(e15, e10)) = op2(h13(e15), h13(e10))) & (h13(op1(e14, e15)) = op2(h13(e14), h13(e15))) & (h13(op1(e14, e14)) = op2(h13(e14), h13(e14))) & (h13(op1(e14, e13)) = op2(h13(e14), h13(e13))) & (h13(op1(e14, e12)) = op2(h13(e14), h13(e12))) & (h13(op1(e14, e11)) = op2(h13(e14), h13(e11))) & (h13(op1(e14, e10)) = op2(h13(e14), h13(e10))) & (h13(op1(e13, e15)) = op2(h13(e13), h13(e15))) & (h13(op1(e13, e14)) = op2(h13(e13), h13(e14))) & (h13(op1(e13, e13)) = op2(h13(e13), h13(e13))) & (h13(op1(e13, e12)) = op2(h13(e13), h13(e12))) & (h13(op1(e13, e11)) = op2(h13(e13), h13(e11))) & (h13(op1(e13, e10)) = op2(h13(e13), h13(e10))) & (h13(op1(e12, e15)) = op2(h13(e12), h13(e15))) & (h13(op1(e12, e14)) = op2(h13(e12), h13(e14))) & (h13(op1(e12, e13)) = op2(h13(e12), h13(e13))) & (h13(op1(e12, e12)) = op2(h13(e12), h13(e12))) & (h13(op1(e12, e11)) = op2(h13(e12), h13(e11))) & (h13(op1(e12, e10)) = op2(h13(e12), h13(e10))) & (h13(op1(e11, e15)) = op2(h13(e11), h13(e15))) & (h13(op1(e11, e14)) = op2(h13(e11), h13(e14))) & (h13(op1(e11, e13)) = op2(h13(e11), h13(e13))) & (h13(op1(e11, e12)) = op2(h13(e11), h13(e12))) & (h13(op1(e11, e11)) = op2(h13(e11), h13(e11))) & (h13(op1(e11, e10)) = op2(h13(e11), h13(e10))) & (h13(op1(e10, e15)) = op2(h13(e10), h13(e15))) & (h13(op1(e10, e14)) = op2(h13(e10), h13(e14))) & (h13(op1(e10, e13)) = op2(h13(e10), h13(e13))) & (h13(op1(e10, e12)) = op2(h13(e10), h13(e12))) & (h13(op1(e10, e11)) = op2(h13(e10), h13(e11))) & (h13(op1(e10, e10)) = op2(h13(e10), h13(e10)))) | (((e25 = h12(e15)) | (e25 = h12(e14)) | (e25 = h12(e13)) | (e25 = h12(e12)) | (e25 = h12(e11)) | (e25 = h12(e10))) & ((e24 = h12(e15)) | (e24 = h12(e14)) | (e24 = h12(e13)) | (e24 = h12(e12)) | (e24 = h12(e11)) | (e24 = h12(e10))) & ((e23 = h12(e15)) | (e23 = h12(e14)) | (e23 = h12(e13)) | (e23 = h12(e12)) | (e23 = h12(e11)) | (e23 = h12(e10))) & ((e22 = h12(e15)) | (e22 = h12(e14)) | (e22 = h12(e13)) | (e22 = h12(e12)) | (e22 = h12(e11)) | (e22 = h12(e10))) & ((e21 = h12(e15)) | (e21 = h12(e14)) | (e21 = h12(e13)) | (e21 = h12(e12)) | (e21 = h12(e11)) | (e21 = h12(e10))) & ((e20 = h12(e15)) | (e20 = h12(e14)) | (e20 = h12(e13)) | (e20 = h12(e12)) | (e20 = h12(e11)) | (e20 = h12(e10))) & (h12(op1(e15, e15)) = op2(h12(e15), h12(e15))) & (h12(op1(e15, e14)) = op2(h12(e15), h12(e14))) & (h12(op1(e15, e13)) = op2(h12(e15), h12(e13))) & (h12(op1(e15, e12)) = op2(h12(e15), h12(e12))) & (h12(op1(e15, e11)) = op2(h12(e15), h12(e11))) & (h12(op1(e15, e10)) = op2(h12(e15), h12(e10))) & (h12(op1(e14, e15)) = op2(h12(e14), h12(e15))) & (h12(op1(e14, e14)) = op2(h12(e14), h12(e14))) & (h12(op1(e14, e13)) = op2(h12(e14), h12(e13))) & (h12(op1(e14, e12)) = op2(h12(e14), h12(e12))) & (h12(op1(e14, e11)) = op2(h12(e14), h12(e11))) & (h12(op1(e14, e10)) = op2(h12(e14), h12(e10))) & (h12(op1(e13, e15)) = op2(h12(e13), h12(e15))) & (h12(op1(e13, e14)) = op2(h12(e13), h12(e14))) & (h12(op1(e13, e13)) = op2(h12(e13), h12(e13))) & (h12(op1(e13, e12)) = op2(h12(e13), h12(e12))) & (h12(op1(e13, e11)) = op2(h12(e13), h12(e11))) & (h12(op1(e13, e10)) = op2(h12(e13), h12(e10))) & (h12(op1(e12, e15)) = op2(h12(e12), h12(e15))) & (h12(op1(e12, e14)) = op2(h12(e12), h12(e14))) & (h12(op1(e12, e13)) = op2(h12(e12), h12(e13))) & (h12(op1(e12, e12)) = op2(h12(e12), h12(e12))) & (h12(op1(e12, e11)) = op2(h12(e12), h12(e11))) & (h12(op1(e12, e10)) = op2(h12(e12), h12(e10))) & (h12(op1(e11, e15)) = op2(h12(e11), h12(e15))) & (h12(op1(e11, e14)) = op2(h12(e11), h12(e14))) & (h12(op1(e11, e13)) = op2(h12(e11), h12(e13))) & (h12(op1(e11, e12)) = op2(h12(e11), h12(e12))) & (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) & (h12(op1(e11, e10)) = op2(h12(e11), h12(e10))) & (h12(op1(e10, e15)) = op2(h12(e10), h12(e15))) & (h12(op1(e10, e14)) = op2(h12(e10), h12(e14))) & (h12(op1(e10, e13)) = op2(h12(e10), h12(e13))) & (h12(op1(e10, e12)) = op2(h12(e10), h12(e12))) & (h12(op1(e10, e11)) = op2(h12(e10), h12(e11))) & (h12(op1(e10, e10)) = op2(h12(e10), h12(e10)))) | (((e25 = h11(e15)) | (e25 = h11(e14)) | (e25 = h11(e13)) | (e25 = h11(e12)) | (e25 = h11(e11)) | (e25 = h11(e10))) & ((e24 = h11(e15)) | (e24 = h11(e14)) | (e24 = h11(e13)) | (e24 = h11(e12)) | (e24 = h11(e11)) | (e24 = h11(e10))) & ((e23 = h11(e15)) | (e23 = h11(e14)) | (e23 = h11(e13)) | (e23 = h11(e12)) | (e23 = h11(e11)) | (e23 = h11(e10))) & ((e22 = h11(e15)) | (e22 = h11(e14)) | (e22 = h11(e13)) | (e22 = h11(e12)) | (e22 = h11(e11)) | (e22 = h11(e10))) & ((e21 = h11(e15)) | (e21 = h11(e14)) | (e21 = h11(e13)) | (e21 = h11(e12)) | (e21 = h11(e11)) | (e21 = h11(e10))) & ((e20 = h11(e15)) | (e20 = h11(e14)) | (e20 = h11(e13)) | (e20 = h11(e12)) | (e20 = h11(e11)) | (e20 = h11(e10))) & (h11(op1(e15, e15)) = op2(h11(e15), h11(e15))) & (h11(op1(e15, e14)) = op2(h11(e15), h11(e14))) & (h11(op1(e15, e13)) = op2(h11(e15), h11(e13))) & (h11(op1(e15, e12)) = op2(h11(e15), h11(e12))) & (h11(op1(e15, e11)) = op2(h11(e15), h11(e11))) & (h11(op1(e15, e10)) = op2(h11(e15), h11(e10))) & (h11(op1(e14, e15)) = op2(h11(e14), h11(e15))) & (h11(op1(e14, e14)) = op2(h11(e14), h11(e14))) & (h11(op1(e14, e13)) = op2(h11(e14), h11(e13))) & (h11(op1(e14, e12)) = op2(h11(e14), h11(e12))) & (h11(op1(e14, e11)) = op2(h11(e14), h11(e11))) & (h11(op1(e14, e10)) = op2(h11(e14), h11(e10))) & (h11(op1(e13, e15)) = op2(h11(e13), h11(e15))) & (h11(op1(e13, e14)) = op2(h11(e13), h11(e14))) & (h11(op1(e13, e13)) = op2(h11(e13), h11(e13))) & (h11(op1(e13, e12)) = op2(h11(e13), h11(e12))) & (h11(op1(e13, e11)) = op2(h11(e13), h11(e11))) & (h11(op1(e13, e10)) = op2(h11(e13), h11(e10))) & (h11(op1(e12, e15)) = op2(h11(e12), h11(e15))) & (h11(op1(e12, e14)) = op2(h11(e12), h11(e14))) & (h11(op1(e12, e13)) = op2(h11(e12), h11(e13))) & (h11(op1(e12, e12)) = op2(h11(e12), h11(e12))) & (h11(op1(e12, e11)) = op2(h11(e12), h11(e11))) & (h11(op1(e12, e10)) = op2(h11(e12), h11(e10))) & (h11(op1(e11, e15)) = op2(h11(e11), h11(e15))) & (h11(op1(e11, e14)) = op2(h11(e11), h11(e14))) & (h11(op1(e11, e13)) = op2(h11(e11), h11(e13))) & (h11(op1(e11, e12)) = op2(h11(e11), h11(e12))) & (h11(op1(e11, e11)) = op2(h11(e11), h11(e11))) & (h11(op1(e11, e10)) = op2(h11(e11), h11(e10))) & (h11(op1(e10, e15)) = op2(h11(e10), h11(e15))) & (h11(op1(e10, e14)) = op2(h11(e10), h11(e14))) & (h11(op1(e10, e13)) = op2(h11(e10), h11(e13))) & (h11(op1(e10, e12)) = op2(h11(e10), h11(e12))) & (h11(op1(e10, e11)) = op2(h11(e10), h11(e11))) & (h11(op1(e10, e10)) = op2(h11(e10), h11(e10)))) | (((e25 = h10(e15)) | (e25 = h10(e14)) | (e25 = h10(e13)) | (e25 = h10(e12)) | (e25 = h10(e11)) | (e25 = h10(e10))) & ((e24 = h10(e15)) | (e24 = h10(e14)) | (e24 = h10(e13)) | (e24 = h10(e12)) | (e24 = h10(e11)) | (e24 = h10(e10))) & ((e23 = h10(e15)) | (e23 = h10(e14)) | (e23 = h10(e13)) | (e23 = h10(e12)) | (e23 = h10(e11)) | (e23 = h10(e10))) & ((e22 = h10(e15)) | (e22 = h10(e14)) | (e22 = h10(e13)) | (e22 = h10(e12)) | (e22 = h10(e11)) | (e22 = h10(e10))) & ((e21 = h10(e15)) | (e21 = h10(e14)) | (e21 = h10(e13)) | (e21 = h10(e12)) | (e21 = h10(e11)) | (e21 = h10(e10))) & ((e20 = h10(e15)) | (e20 = h10(e14)) | (e20 = h10(e13)) | (e20 = h10(e12)) | (e20 = h10(e11)) | (e20 = h10(e10))) & (h10(op1(e15, e15)) = op2(h10(e15), h10(e15))) & (h10(op1(e15, e14)) = op2(h10(e15), h10(e14))) & (h10(op1(e15, e13)) = op2(h10(e15), h10(e13))) & (h10(op1(e15, e12)) = op2(h10(e15), h10(e12))) & (h10(op1(e15, e11)) = op2(h10(e15), h10(e11))) & (h10(op1(e15, e10)) = op2(h10(e15), h10(e10))) & (h10(op1(e14, e15)) = op2(h10(e14), h10(e15))) & (h10(op1(e14, e14)) = op2(h10(e14), h10(e14))) & (h10(op1(e14, e13)) = op2(h10(e14), h10(e13))) & (h10(op1(e14, e12)) = op2(h10(e14), h10(e12))) & (h10(op1(e14, e11)) = op2(h10(e14), h10(e11))) & (h10(op1(e14, e10)) = op2(h10(e14), h10(e10))) & (h10(op1(e13, e15)) = op2(h10(e13), h10(e15))) & (h10(op1(e13, e14)) = op2(h10(e13), h10(e14))) & (h10(op1(e13, e13)) = op2(h10(e13), h10(e13))) & (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) & (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) & (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) & (h10(op1(e12, e15)) = op2(h10(e12), h10(e15))) & (h10(op1(e12, e14)) = op2(h10(e12), h10(e14))) & (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) & (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) & (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) & (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) & (h10(op1(e11, e15)) = op2(h10(e11), h10(e15))) & (h10(op1(e11, e14)) = op2(h10(e11), h10(e14))) & (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) & (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) & (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) & (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) & (h10(op1(e10, e15)) = op2(h10(e10), h10(e15))) & (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) & (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) & (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) & (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) & (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))) | (((e25 = h9(e15)) | (e25 = h9(e14)) | (e25 = h9(e13)) | (e25 = h9(e12)) | (e25 = h9(e11)) | (e25 = h9(e10))) & ((e24 = h9(e15)) | (e24 = h9(e14)) | (e24 = h9(e13)) | (e24 = h9(e12)) | (e24 = h9(e11)) | (e24 = h9(e10))) & ((e23 = h9(e15)) | (e23 = h9(e14)) | (e23 = h9(e13)) | (e23 = h9(e12)) | (e23 = h9(e11)) | (e23 = h9(e10))) & ((e22 = h9(e15)) | (e22 = h9(e14)) | (e22 = h9(e13)) | (e22 = h9(e12)) | (e22 = h9(e11)) | (e22 = h9(e10))) & ((e21 = h9(e15)) | (e21 = h9(e14)) | (e21 = h9(e13)) | (e21 = h9(e12)) | (e21 = h9(e11)) | (e21 = h9(e10))) & ((e20 = h9(e15)) | (e20 = h9(e14)) | (e20 = h9(e13)) | (e20 = h9(e12)) | (e20 = h9(e11)) | (e20 = h9(e10))) & (h9(op1(e15, e15)) = op2(h9(e15), h9(e15))) & (h9(op1(e15, e14)) = op2(h9(e15), h9(e14))) & (h9(op1(e15, e13)) = op2(h9(e15), h9(e13))) & (h9(op1(e15, e12)) = op2(h9(e15), h9(e12))) & (h9(op1(e15, e11)) = op2(h9(e15), h9(e11))) & (h9(op1(e15, e10)) = op2(h9(e15), h9(e10))) & (h9(op1(e14, e15)) = op2(h9(e14), h9(e15))) & (h9(op1(e14, e14)) = op2(h9(e14), h9(e14))) & (h9(op1(e14, e13)) = op2(h9(e14), h9(e13))) & (h9(op1(e14, e12)) = op2(h9(e14), h9(e12))) & (h9(op1(e14, e11)) = op2(h9(e14), h9(e11))) & (h9(op1(e14, e10)) = op2(h9(e14), h9(e10))) & (h9(op1(e13, e15)) = op2(h9(e13), h9(e15))) & (h9(op1(e13, e14)) = op2(h9(e13), h9(e14))) & (h9(op1(e13, e13)) = op2(h9(e13), h9(e13))) & (h9(op1(e13, e12)) = op2(h9(e13), h9(e12))) & (h9(op1(e13, e11)) = op2(h9(e13), h9(e11))) & (h9(op1(e13, e10)) = op2(h9(e13), h9(e10))) & (h9(op1(e12, e15)) = op2(h9(e12), h9(e15))) & (h9(op1(e12, e14)) = op2(h9(e12), h9(e14))) & (h9(op1(e12, e13)) = op2(h9(e12), h9(e13))) & (h9(op1(e12, e12)) = op2(h9(e12), h9(e12))) & (h9(op1(e12, e11)) = op2(h9(e12), h9(e11))) & (h9(op1(e12, e10)) = op2(h9(e12), h9(e10))) & (h9(op1(e11, e15)) = op2(h9(e11), h9(e15))) & (h9(op1(e11, e14)) = op2(h9(e11), h9(e14))) & (h9(op1(e11, e13)) = op2(h9(e11), h9(e13))) & (h9(op1(e11, e12)) = op2(h9(e11), h9(e12))) & (h9(op1(e11, e11)) = op2(h9(e11), h9(e11))) & (h9(op1(e11, e10)) = op2(h9(e11), h9(e10))) & (h9(op1(e10, e15)) = op2(h9(e10), h9(e15))) & (h9(op1(e10, e14)) = op2(h9(e10), h9(e14))) & (h9(op1(e10, e13)) = op2(h9(e10), h9(e13))) & (h9(op1(e10, e12)) = op2(h9(e10), h9(e12))) & (h9(op1(e10, e11)) = op2(h9(e10), h9(e11))) & (h9(op1(e10, e10)) = op2(h9(e10), h9(e10)))) | (((e25 = h8(e15)) | (e25 = h8(e14)) | (e25 = h8(e13)) | (e25 = h8(e12)) | (e25 = h8(e11)) | (e25 = h8(e10))) & ((e24 = h8(e15)) | (e24 = h8(e14)) | (e24 = h8(e13)) | (e24 = h8(e12)) | (e24 = h8(e11)) | (e24 = h8(e10))) & ((e23 = h8(e15)) | (e23 = h8(e14)) | (e23 = h8(e13)) | (e23 = h8(e12)) | (e23 = h8(e11)) | (e23 = h8(e10))) & ((e22 = h8(e15)) | (e22 = h8(e14)) | (e22 = h8(e13)) | (e22 = h8(e12)) | (e22 = h8(e11)) | (e22 = h8(e10))) & ((e21 = h8(e15)) | (e21 = h8(e14)) | (e21 = h8(e13)) | (e21 = h8(e12)) | (e21 = h8(e11)) | (e21 = h8(e10))) & ((e20 = h8(e15)) | (e20 = h8(e14)) | (e20 = h8(e13)) | (e20 = h8(e12)) | (e20 = h8(e11)) | (e20 = h8(e10))) & (h8(op1(e15, e15)) = op2(h8(e15), h8(e15))) & (h8(op1(e15, e14)) = op2(h8(e15), h8(e14))) & (h8(op1(e15, e13)) = op2(h8(e15), h8(e13))) & (h8(op1(e15, e12)) = op2(h8(e15), h8(e12))) & (h8(op1(e15, e11)) = op2(h8(e15), h8(e11))) & (h8(op1(e15, e10)) = op2(h8(e15), h8(e10))) & (h8(op1(e14, e15)) = op2(h8(e14), h8(e15))) & (h8(op1(e14, e14)) = op2(h8(e14), h8(e14))) & (h8(op1(e14, e13)) = op2(h8(e14), h8(e13))) & (h8(op1(e14, e12)) = op2(h8(e14), h8(e12))) & (h8(op1(e14, e11)) = op2(h8(e14), h8(e11))) & (h8(op1(e14, e10)) = op2(h8(e14), h8(e10))) & (h8(op1(e13, e15)) = op2(h8(e13), h8(e15))) & (h8(op1(e13, e14)) = op2(h8(e13), h8(e14))) & (h8(op1(e13, e13)) = op2(h8(e13), h8(e13))) & (h8(op1(e13, e12)) = op2(h8(e13), h8(e12))) & (h8(op1(e13, e11)) = op2(h8(e13), h8(e11))) & (h8(op1(e13, e10)) = op2(h8(e13), h8(e10))) & (h8(op1(e12, e15)) = op2(h8(e12), h8(e15))) & (h8(op1(e12, e14)) = op2(h8(e12), h8(e14))) & (h8(op1(e12, e13)) = op2(h8(e12), h8(e13))) & (h8(op1(e12, e12)) = op2(h8(e12), h8(e12))) & (h8(op1(e12, e11)) = op2(h8(e12), h8(e11))) & (h8(op1(e12, e10)) = op2(h8(e12), h8(e10))) & (h8(op1(e11, e15)) = op2(h8(e11), h8(e15))) & (h8(op1(e11, e14)) = op2(h8(e11), h8(e14))) & (h8(op1(e11, e13)) = op2(h8(e11), h8(e13))) & (h8(op1(e11, e12)) = op2(h8(e11), h8(e12))) & (h8(op1(e11, e11)) = op2(h8(e11), h8(e11))) & (h8(op1(e11, e10)) = op2(h8(e11), h8(e10))) & (h8(op1(e10, e15)) = op2(h8(e10), h8(e15))) & (h8(op1(e10, e14)) = op2(h8(e10), h8(e14))) & (h8(op1(e10, e13)) = op2(h8(e10), h8(e13))) & (h8(op1(e10, e12)) = op2(h8(e10), h8(e12))) & (h8(op1(e10, e11)) = op2(h8(e10), h8(e11))) & (h8(op1(e10, e10)) = op2(h8(e10), h8(e10)))) | (((e25 = h7(e15)) | (e25 = h7(e14)) | (e25 = h7(e13)) | (e25 = h7(e12)) | (e25 = h7(e11)) | (e25 = h7(e10))) & ((e24 = h7(e15)) | (e24 = h7(e14)) | (e24 = h7(e13)) | (e24 = h7(e12)) | (e24 = h7(e11)) | (e24 = h7(e10))) & ((e23 = h7(e15)) | (e23 = h7(e14)) | (e23 = h7(e13)) | (e23 = h7(e12)) | (e23 = h7(e11)) | (e23 = h7(e10))) & ((e22 = h7(e15)) | (e22 = h7(e14)) | (e22 = h7(e13)) | (e22 = h7(e12)) | (e22 = h7(e11)) | (e22 = h7(e10))) & ((e21 = h7(e15)) | (e21 = h7(e14)) | (e21 = h7(e13)) | (e21 = h7(e12)) | (e21 = h7(e11)) | (e21 = h7(e10))) & ((e20 = h7(e15)) | (e20 = h7(e14)) | (e20 = h7(e13)) | (e20 = h7(e12)) | (e20 = h7(e11)) | (e20 = h7(e10))) & (h7(op1(e15, e15)) = op2(h7(e15), h7(e15))) & (h7(op1(e15, e14)) = op2(h7(e15), h7(e14))) & (h7(op1(e15, e13)) = op2(h7(e15), h7(e13))) & (h7(op1(e15, e12)) = op2(h7(e15), h7(e12))) & (h7(op1(e15, e11)) = op2(h7(e15), h7(e11))) & (h7(op1(e15, e10)) = op2(h7(e15), h7(e10))) & (h7(op1(e14, e15)) = op2(h7(e14), h7(e15))) & (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) & (h7(op1(e14, e13)) = op2(h7(e14), h7(e13))) & (h7(op1(e14, e12)) = op2(h7(e14), h7(e12))) & (h7(op1(e14, e11)) = op2(h7(e14), h7(e11))) & (h7(op1(e14, e10)) = op2(h7(e14), h7(e10))) & (h7(op1(e13, e15)) = op2(h7(e13), h7(e15))) & (h7(op1(e13, e14)) = op2(h7(e13), h7(e14))) & (h7(op1(e13, e13)) = op2(h7(e13), h7(e13))) & (h7(op1(e13, e12)) = op2(h7(e13), h7(e12))) & (h7(op1(e13, e11)) = op2(h7(e13), h7(e11))) & (h7(op1(e13, e10)) = op2(h7(e13), h7(e10))) & (h7(op1(e12, e15)) = op2(h7(e12), h7(e15))) & (h7(op1(e12, e14)) = op2(h7(e12), h7(e14))) & (h7(op1(e12, e13)) = op2(h7(e12), h7(e13))) & (h7(op1(e12, e12)) = op2(h7(e12), h7(e12))) & (h7(op1(e12, e11)) = op2(h7(e12), h7(e11))) & (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) & (h7(op1(e11, e15)) = op2(h7(e11), h7(e15))) & (h7(op1(e11, e14)) = op2(h7(e11), h7(e14))) & (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) & (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) & (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) & (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) & (h7(op1(e10, e15)) = op2(h7(e10), h7(e15))) & (h7(op1(e10, e14)) = op2(h7(e10), h7(e14))) & (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) & (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) & (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) & (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))) | (((e25 = h6(e15)) | (e25 = h6(e14)) | (e25 = h6(e13)) | (e25 = h6(e12)) | (e25 = h6(e11)) | (e25 = h6(e10))) & ((e24 = h6(e15)) | (e24 = h6(e14)) | (e24 = h6(e13)) | (e24 = h6(e12)) | (e24 = h6(e11)) | (e24 = h6(e10))) & ((e23 = h6(e15)) | (e23 = h6(e14)) | (e23 = h6(e13)) | (e23 = h6(e12)) | (e23 = h6(e11)) | (e23 = h6(e10))) & ((e22 = h6(e15)) | (e22 = h6(e14)) | (e22 = h6(e13)) | (e22 = h6(e12)) | (e22 = h6(e11)) | (e22 = h6(e10))) & ((e21 = h6(e15)) | (e21 = h6(e14)) | (e21 = h6(e13)) | (e21 = h6(e12)) | (e21 = h6(e11)) | (e21 = h6(e10))) & ((e20 = h6(e15)) | (e20 = h6(e14)) | (e20 = h6(e13)) | (e20 = h6(e12)) | (e20 = h6(e11)) | (e20 = h6(e10))) & (h6(op1(e15, e15)) = op2(h6(e15), h6(e15))) & (h6(op1(e15, e14)) = op2(h6(e15), h6(e14))) & (h6(op1(e15, e13)) = op2(h6(e15), h6(e13))) & (h6(op1(e15, e12)) = op2(h6(e15), h6(e12))) & (h6(op1(e15, e11)) = op2(h6(e15), h6(e11))) & (h6(op1(e15, e10)) = op2(h6(e15), h6(e10))) & (h6(op1(e14, e15)) = op2(h6(e14), h6(e15))) & (h6(op1(e14, e14)) = op2(h6(e14), h6(e14))) & (h6(op1(e14, e13)) = op2(h6(e14), h6(e13))) & (h6(op1(e14, e12)) = op2(h6(e14), h6(e12))) & (h6(op1(e14, e11)) = op2(h6(e14), h6(e11))) & (h6(op1(e14, e10)) = op2(h6(e14), h6(e10))) & (h6(op1(e13, e15)) = op2(h6(e13), h6(e15))) & (h6(op1(e13, e14)) = op2(h6(e13), h6(e14))) & (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) & (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) & (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) & (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) & (h6(op1(e12, e15)) = op2(h6(e12), h6(e15))) & (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) & (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) & (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) & (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) & (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) & (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) & (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) & (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) & (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) & (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) & (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) & (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) & (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) & (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) & (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) & (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) & (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))) | (((e25 = h5(e15)) | (e25 = h5(e14)) | (e25 = h5(e13)) | (e25 = h5(e12)) | (e25 = h5(e11)) | (e25 = h5(e10))) & ((e24 = h5(e15)) | (e24 = h5(e14)) | (e24 = h5(e13)) | (e24 = h5(e12)) | (e24 = h5(e11)) | (e24 = h5(e10))) & ((e23 = h5(e15)) | (e23 = h5(e14)) | (e23 = h5(e13)) | (e23 = h5(e12)) | (e23 = h5(e11)) | (e23 = h5(e10))) & ((e22 = h5(e15)) | (e22 = h5(e14)) | (e22 = h5(e13)) | (e22 = h5(e12)) | (e22 = h5(e11)) | (e22 = h5(e10))) & ((e21 = h5(e15)) | (e21 = h5(e14)) | (e21 = h5(e13)) | (e21 = h5(e12)) | (e21 = h5(e11)) | (e21 = h5(e10))) & ((e20 = h5(e15)) | (e20 = h5(e14)) | (e20 = h5(e13)) | (e20 = h5(e12)) | (e20 = h5(e11)) | (e20 = h5(e10))) & (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) & (h5(op1(e15, e14)) = op2(h5(e15), h5(e14))) & (h5(op1(e15, e13)) = op2(h5(e15), h5(e13))) & (h5(op1(e15, e12)) = op2(h5(e15), h5(e12))) & (h5(op1(e15, e11)) = op2(h5(e15), h5(e11))) & (h5(op1(e15, e10)) = op2(h5(e15), h5(e10))) & (h5(op1(e14, e15)) = op2(h5(e14), h5(e15))) & (h5(op1(e14, e14)) = op2(h5(e14), h5(e14))) & (h5(op1(e14, e13)) = op2(h5(e14), h5(e13))) & (h5(op1(e14, e12)) = op2(h5(e14), h5(e12))) & (h5(op1(e14, e11)) = op2(h5(e14), h5(e11))) & (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) & (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) & (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) & (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) & (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) & (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) & (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) & (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) & (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) & (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) & (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) & (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) & (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) & (h5(op1(e11, e15)) = op2(h5(e11), h5(e15))) & (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) & (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) & (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) & (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) & (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) & (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) & (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) & (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) & (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) & (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) & (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))) | (((e25 = h4(e15)) | (e25 = h4(e14)) | (e25 = h4(e13)) | (e25 = h4(e12)) | (e25 = h4(e11)) | (e25 = h4(e10))) & ((e24 = h4(e15)) | (e24 = h4(e14)) | (e24 = h4(e13)) | (e24 = h4(e12)) | (e24 = h4(e11)) | (e24 = h4(e10))) & ((e23 = h4(e15)) | (e23 = h4(e14)) | (e23 = h4(e13)) | (e23 = h4(e12)) | (e23 = h4(e11)) | (e23 = h4(e10))) & ((e22 = h4(e15)) | (e22 = h4(e14)) | (e22 = h4(e13)) | (e22 = h4(e12)) | (e22 = h4(e11)) | (e22 = h4(e10))) & ((e21 = h4(e15)) | (e21 = h4(e14)) | (e21 = h4(e13)) | (e21 = h4(e12)) | (e21 = h4(e11)) | (e21 = h4(e10))) & ((e20 = h4(e15)) | (e20 = h4(e14)) | (e20 = h4(e13)) | (e20 = h4(e12)) | (e20 = h4(e11)) | (e20 = h4(e10))) & (h4(op1(e15, e15)) = op2(h4(e15), h4(e15))) & (h4(op1(e15, e14)) = op2(h4(e15), h4(e14))) & (h4(op1(e15, e13)) = op2(h4(e15), h4(e13))) & (h4(op1(e15, e12)) = op2(h4(e15), h4(e12))) & (h4(op1(e15, e11)) = op2(h4(e15), h4(e11))) & (h4(op1(e15, e10)) = op2(h4(e15), h4(e10))) & (h4(op1(e14, e15)) = op2(h4(e14), h4(e15))) & (h4(op1(e14, e14)) = op2(h4(e14), h4(e14))) & (h4(op1(e14, e13)) = op2(h4(e14), h4(e13))) & (h4(op1(e14, e12)) = op2(h4(e14), h4(e12))) & (h4(op1(e14, e11)) = op2(h4(e14), h4(e11))) & (h4(op1(e14, e10)) = op2(h4(e14), h4(e10))) & (h4(op1(e13, e15)) = op2(h4(e13), h4(e15))) & (h4(op1(e13, e14)) = op2(h4(e13), h4(e14))) & (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) & (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) & (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) & (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) & (h4(op1(e12, e15)) = op2(h4(e12), h4(e15))) & (h4(op1(e12, e14)) = op2(h4(e12), h4(e14))) & (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) & (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) & (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) & (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) & (h4(op1(e11, e15)) = op2(h4(e11), h4(e15))) & (h4(op1(e11, e14)) = op2(h4(e11), h4(e14))) & (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) & (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) & (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) & (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) & (h4(op1(e10, e15)) = op2(h4(e10), h4(e15))) & (h4(op1(e10, e14)) = op2(h4(e10), h4(e14))) & (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) & (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) & (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) & (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) | (((e25 = h3(e15)) | (e25 = h3(e14)) | (e25 = h3(e13)) | (e25 = h3(e12)) | (e25 = h3(e11)) | (e25 = h3(e10))) & ((e24 = h3(e15)) | (e24 = h3(e14)) | (e24 = h3(e13)) | (e24 = h3(e12)) | (e24 = h3(e11)) | (e24 = h3(e10))) & ((e23 = h3(e15)) | (e23 = h3(e14)) | (e23 = h3(e13)) | (e23 = h3(e12)) | (e23 = h3(e11)) | (e23 = h3(e10))) & ((e22 = h3(e15)) | (e22 = h3(e14)) | (e22 = h3(e13)) | (e22 = h3(e12)) | (e22 = h3(e11)) | (e22 = h3(e10))) & ((e21 = h3(e15)) | (e21 = h3(e14)) | (e21 = h3(e13)) | (e21 = h3(e12)) | (e21 = h3(e11)) | (e21 = h3(e10))) & ((e20 = h3(e15)) | (e20 = h3(e14)) | (e20 = h3(e13)) | (e20 = h3(e12)) | (e20 = h3(e11)) | (e20 = h3(e10))) & (h3(op1(e15, e15)) = op2(h3(e15), h3(e15))) & (h3(op1(e15, e14)) = op2(h3(e15), h3(e14))) & (h3(op1(e15, e13)) = op2(h3(e15), h3(e13))) & (h3(op1(e15, e12)) = op2(h3(e15), h3(e12))) & (h3(op1(e15, e11)) = op2(h3(e15), h3(e11))) & (h3(op1(e15, e10)) = op2(h3(e15), h3(e10))) & (h3(op1(e14, e15)) = op2(h3(e14), h3(e15))) & (h3(op1(e14, e14)) = op2(h3(e14), h3(e14))) & (h3(op1(e14, e13)) = op2(h3(e14), h3(e13))) & (h3(op1(e14, e12)) = op2(h3(e14), h3(e12))) & (h3(op1(e14, e11)) = op2(h3(e14), h3(e11))) & (h3(op1(e14, e10)) = op2(h3(e14), h3(e10))) & (h3(op1(e13, e15)) = op2(h3(e13), h3(e15))) & (h3(op1(e13, e14)) = op2(h3(e13), h3(e14))) & (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) & (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) & (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) & (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) & (h3(op1(e12, e15)) = op2(h3(e12), h3(e15))) & (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) & (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) & (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) & (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) & (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) & (h3(op1(e11, e15)) = op2(h3(e11), h3(e15))) & (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) & (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) & (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) & (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) & (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) & (h3(op1(e10, e15)) = op2(h3(e10), h3(e15))) & (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) & (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) & (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) & (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) & (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) | (((e25 = h2(e15)) | (e25 = h2(e14)) | (e25 = h2(e13)) | (e25 = h2(e12)) | (e25 = h2(e11)) | (e25 = h2(e10))) & ((e24 = h2(e15)) | (e24 = h2(e14)) | (e24 = h2(e13)) | (e24 = h2(e12)) | (e24 = h2(e11)) | (e24 = h2(e10))) & ((e23 = h2(e15)) | (e23 = h2(e14)) | (e23 = h2(e13)) | (e23 = h2(e12)) | (e23 = h2(e11)) | (e23 = h2(e10))) & ((e22 = h2(e15)) | (e22 = h2(e14)) | (e22 = h2(e13)) | (e22 = h2(e12)) | (e22 = h2(e11)) | (e22 = h2(e10))) & ((e21 = h2(e15)) | (e21 = h2(e14)) | (e21 = h2(e13)) | (e21 = h2(e12)) | (e21 = h2(e11)) | (e21 = h2(e10))) & ((e20 = h2(e15)) | (e20 = h2(e14)) | (e20 = h2(e13)) | (e20 = h2(e12)) | (e20 = h2(e11)) | (e20 = h2(e10))) & (h2(op1(e15, e15)) = op2(h2(e15), h2(e15))) & (h2(op1(e15, e14)) = op2(h2(e15), h2(e14))) & (h2(op1(e15, e13)) = op2(h2(e15), h2(e13))) & (h2(op1(e15, e12)) = op2(h2(e15), h2(e12))) & (h2(op1(e15, e11)) = op2(h2(e15), h2(e11))) & (h2(op1(e15, e10)) = op2(h2(e15), h2(e10))) & (h2(op1(e14, e15)) = op2(h2(e14), h2(e15))) & (h2(op1(e14, e14)) = op2(h2(e14), h2(e14))) & (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) & (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) & (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) & (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) & (h2(op1(e13, e15)) = op2(h2(e13), h2(e15))) & (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) & (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) & (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) & (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) & (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) & (h2(op1(e12, e15)) = op2(h2(e12), h2(e15))) & (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) & (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) & (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) & (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) & (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) & (h2(op1(e11, e15)) = op2(h2(e11), h2(e15))) & (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) & (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) & (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) & (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) & (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) & (h2(op1(e10, e15)) = op2(h2(e10), h2(e15))) & (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) & (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) & (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) & (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) & (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) | (((e25 = h1(e15)) | (e25 = h1(e14)) | (e25 = h1(e13)) | (e25 = h1(e12)) | (e25 = h1(e11)) | (e25 = h1(e10))) & ((e24 = h1(e15)) | (e24 = h1(e14)) | (e24 = h1(e13)) | (e24 = h1(e12)) | (e24 = h1(e11)) | (e24 = h1(e10))) & ((e23 = h1(e15)) | (e23 = h1(e14)) | (e23 = h1(e13)) | (e23 = h1(e12)) | (e23 = h1(e11)) | (e23 = h1(e10))) & ((e22 = h1(e15)) | (e22 = h1(e14)) | (e22 = h1(e13)) | (e22 = h1(e12)) | (e22 = h1(e11)) | (e22 = h1(e10))) & ((e21 = h1(e15)) | (e21 = h1(e14)) | (e21 = h1(e13)) | (e21 = h1(e12)) | (e21 = h1(e11)) | (e21 = h1(e10))) & ((e20 = h1(e15)) | (e20 = h1(e14)) | (e20 = h1(e13)) | (e20 = h1(e12)) | (e20 = h1(e11)) | (e20 = h1(e10))) & (h1(op1(e15, e15)) = op2(h1(e15), h1(e15))) & (h1(op1(e15, e14)) = op2(h1(e15), h1(e14))) & (h1(op1(e15, e13)) = op2(h1(e15), h1(e13))) & (h1(op1(e15, e12)) = op2(h1(e15), h1(e12))) & (h1(op1(e15, e11)) = op2(h1(e15), h1(e11))) & (h1(op1(e15, e10)) = op2(h1(e15), h1(e10))) & (h1(op1(e14, e15)) = op2(h1(e14), h1(e15))) & (h1(op1(e14, e14)) = op2(h1(e14), h1(e14))) & (h1(op1(e14, e13)) = op2(h1(e14), h1(e13))) & (h1(op1(e14, e12)) = op2(h1(e14), h1(e12))) & (h1(op1(e14, e11)) = op2(h1(e14), h1(e11))) & (h1(op1(e14, e10)) = op2(h1(e14), h1(e10))) & (h1(op1(e13, e15)) = op2(h1(e13), h1(e15))) & (h1(op1(e13, e14)) = op2(h1(e13), h1(e14))) & (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) & (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) & (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) & (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) & (h1(op1(e12, e15)) = op2(h1(e12), h1(e15))) & (h1(op1(e12, e14)) = op2(h1(e12), h1(e14))) & (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) & (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) & (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) & (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) & (h1(op1(e11, e15)) = op2(h1(e11), h1(e15))) & (h1(op1(e11, e14)) = op2(h1(e11), h1(e14))) & (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) & (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) & (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) & (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) & (h1(op1(e10, e15)) = op2(h1(e10), h1(e15))) & (h1(op1(e10, e14)) = op2(h1(e10), h1(e14))) & (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) & (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) & (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) & (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG027+1.p', co1)).
fof(f8039, plain, ~ spl150_959, inference(avatar_split_clause, [], [f8038, f8024])).
fof(f8038, plain, ~ sP45, inference(subsumption_resolution, [], [f2286, f5365])).
fof(f2286, plain, (~ (e20 = h10(e10)) | ~ sP45), inference(cnf_transformation, [], [f318])).
fof(f318, plain, ((~ (e20 = h10(e15)) & ~ (e20 = h10(e14)) & ~ (e20 = h10(e13)) & ~ (e20 = h10(e12)) & ~ (e20 = h10(e11)) & ~ (e20 = h10(e10))) | ~ sP45), inference(nnf_transformation, [], [f108])).
fof(f8019, plain, ~ spl150_957, inference(avatar_split_clause, [], [f8018, f8009])).
fof(f8018, plain, ~ sP46, inference(subsumption_resolution, [], [f2281, f1536])).
fof(f2281, plain, (~ (e21 = h10(e11)) | ~ sP46), inference(cnf_transformation, [], [f317])).
fof(f317, plain, ((~ (e21 = h10(e15)) & ~ (e21 = h10(e14)) & ~ (e21 = h10(e13)) & ~ (e21 = h10(e12)) & ~ (e21 = h10(e11)) & ~ (e21 = h10(e10))) | ~ sP46), inference(nnf_transformation, [], [f109])).
fof(f8003, plain, (~ spl150_955 | ~ spl150_656), inference(avatar_split_clause, [], [f2276, f6398, f7995])).
fof(f2276, plain, (~ (e22 = h10(e12)) | ~ sP47), inference(cnf_transformation, [], [f316])).
fof(f316, plain, ((~ (e22 = h10(e15)) & ~ (e22 = h10(e14)) & ~ (e22 = h10(e13)) & ~ (e22 = h10(e12)) & ~ (e22 = h10(e11)) & ~ (e22 = h10(e10))) | ~ sP47), inference(nnf_transformation, [], [f110])).
fof(f7988, plain, (~ spl150_953 | ~ spl150_954), inference(avatar_split_clause, [], [f2271, f7985, f7981])).
fof(f2271, plain, (~ (e23 = h10(e13)) | ~ sP48), inference(cnf_transformation, [], [f315])).
fof(f315, plain, ((~ (e23 = h10(e15)) & ~ (e23 = h10(e14)) & ~ (e23 = h10(e13)) & ~ (e23 = h10(e12)) & ~ (e23 = h10(e11)) & ~ (e23 = h10(e10))) | ~ sP48), inference(nnf_transformation, [], [f111])).
fof(f5084, plain, spl150_267, inference(avatar_split_clause, [], [f5083, f3855])).
fof(f5083, plain, (e22 = op2(e25, e24)), inference(forward_demodulation, [], [f1479, f1481])).
fof(f1479, plain, (e22 = op2(e25, op2(op2(e25, e21), op2(e25, e21)))), inference(cnf_transformation, [], [f25])).
fof(f5082, plain, spl150_286, inference(avatar_split_clause, [], [f1480, f3934])).
fof(f1480, plain, (e23 = op2(e25, e21)), inference(cnf_transformation, [], [f25])).
fof(f5080, plain, spl150_9, inference(avatar_split_clause, [], [f5079, f2780])).
fof(f5079, plain, (e12 = op1(e15, e14)), inference(forward_demodulation, [], [f1475, f1477])).
fof(f1475, plain, (e12 = op1(e15, op1(op1(e15, e11), op1(e15, e11)))), inference(cnf_transformation, [], [f24])).
fof(f5078, plain, spl150_28, inference(avatar_split_clause, [], [f1476, f2859])).
fof(f1476, plain, (e13 = op1(e15, e11)), inference(cnf_transformation, [], [f24])).
fof(f4944, plain, (~ spl150_500 | spl150_507), inference(avatar_split_clause, [], [f1008, f4855, f4826])).
fof(f4826, plain, (spl150_500 <=> (e21 = inv2(e22))), introduced(avatar_definition, [new_symbols(naming, [spl150_500])])).
fof(f1008, plain, ((e22 = inv2(e21)) | ~ (e21 = inv2(e22))), inference(cnf_transformation, [], [f59])).
fof(f59, plain, (((e25 = inv2(e25)) | ~ (e25 = inv2(e25))) & ((e25 = inv2(e24)) | ~ (e24 = inv2(e25))) & ((e25 = inv2(e23)) | ~ (e23 = inv2(e25))) & ((e25 = inv2(e22)) | ~ (e22 = inv2(e25))) & ((e25 = inv2(e21)) | ~ (e21 = inv2(e25))) & ((e25 = inv2(e20)) | ~ (e20 = inv2(e25))) & ((e24 = inv2(e25)) | ~ (e25 = inv2(e24))) & ((e24 = inv2(e24)) | ~ (e24 = inv2(e24))) & ((e24 = inv2(e23)) | ~ (e23 = inv2(e24))) & ((e24 = inv2(e22)) | ~ (e22 = inv2(e24))) & ((e24 = inv2(e21)) | ~ (e21 = inv2(e24))) & ((e24 = inv2(e20)) | ~ (e20 = inv2(e24))) & ((e23 = inv2(e25)) | ~ (e25 = inv2(e23))) & ((e23 = inv2(e24)) | ~ (e24 = inv2(e23))) & ((e23 = inv2(e23)) | ~ (e23 = inv2(e23))) & ((e23 = inv2(e22)) | ~ (e22 = inv2(e23))) & ((e23 = inv2(e21)) | ~ (e21 = inv2(e23))) & ((e23 = inv2(e20)) | ~ (e20 = inv2(e23))) & ((e22 = inv2(e25)) | ~ (e25 = inv2(e22))) & ((e22 = inv2(e24)) | ~ (e24 = inv2(e22))) & ((e22 = inv2(e23)) | ~ (e23 = inv2(e22))) & ((e22 = inv2(e22)) | ~ (e22 = inv2(e22))) & ((e22 = inv2(e21)) | ~ (e21 = inv2(e22))) & ((e22 = inv2(e20)) | ~ (e20 = inv2(e22))) & ((e21 = inv2(e25)) | ~ (e25 = inv2(e21))) & ((e21 = inv2(e24)) | ~ (e24 = inv2(e21))) & ((e21 = inv2(e23)) | ~ (e23 = inv2(e21))) & ((e21 = inv2(e22)) | ~ (e22 = inv2(e21))) & ((e21 = inv2(e21)) | ~ (e21 = inv2(e21))) & ((e21 = inv2(e20)) | ~ (e20 = inv2(e21))) & ((e20 = inv2(e25)) | ~ (e25 = inv2(e20))) & ((e20 = inv2(e24)) | ~ (e24 = inv2(e20))) & ((e20 = inv2(e23)) | ~ (e23 = inv2(e20))) & ((e20 = inv2(e22)) | ~ (e22 = inv2(e20))) & ((e20 = inv2(e21)) | ~ (e21 = inv2(e20))) & ((e20 = inv2(e20)) | ~ (e20 = inv2(e20)))), inference(ennf_transformation, [], [f15])).
fof(f15, plain, (((e25 = inv2(e25)) => (e25 = inv2(e25))) & ((e24 = inv2(e25)) => (e25 = inv2(e24))) & ((e23 = inv2(e25)) => (e25 = inv2(e23))) & ((e22 = inv2(e25)) => (e25 = inv2(e22))) & ((e21 = inv2(e25)) => (e25 = inv2(e21))) & ((e20 = inv2(e25)) => (e25 = inv2(e20))) & ((e25 = inv2(e24)) => (e24 = inv2(e25))) & ((e24 = inv2(e24)) => (e24 = inv2(e24))) & ((e23 = inv2(e24)) => (e24 = inv2(e23))) & ((e22 = inv2(e24)) => (e24 = inv2(e22))) & ((e21 = inv2(e24)) => (e24 = inv2(e21))) & ((e20 = inv2(e24)) => (e24 = inv2(e20))) & ((e25 = inv2(e23)) => (e23 = inv2(e25))) & ((e24 = inv2(e23)) => (e23 = inv2(e24))) & ((e23 = inv2(e23)) => (e23 = inv2(e23))) & ((e22 = inv2(e23)) => (e23 = inv2(e22))) & ((e21 = inv2(e23)) => (e23 = inv2(e21))) & ((e20 = inv2(e23)) => (e23 = inv2(e20))) & ((e25 = inv2(e22)) => (e22 = inv2(e25))) & ((e24 = inv2(e22)) => (e22 = inv2(e24))) & ((e23 = inv2(e22)) => (e22 = inv2(e23))) & ((e22 = inv2(e22)) => (e22 = inv2(e22))) & ((e21 = inv2(e22)) => (e22 = inv2(e21))) & ((e20 = inv2(e22)) => (e22 = inv2(e20))) & ((e25 = inv2(e21)) => (e21 = inv2(e25))) & ((e24 = inv2(e21)) => (e21 = inv2(e24))) & ((e23 = inv2(e21)) => (e21 = inv2(e23))) & ((e22 = inv2(e21)) => (e21 = inv2(e22))) & ((e21 = inv2(e21)) => (e21 = inv2(e21))) & ((e20 = inv2(e21)) => (e21 = inv2(e20))) & ((e25 = inv2(e20)) => (e20 = inv2(e25))) & ((e24 = inv2(e20)) => (e20 = inv2(e24))) & ((e23 = inv2(e20)) => (e20 = inv2(e23))) & ((e22 = inv2(e20)) => (e20 = inv2(e22))) & ((e21 = inv2(e20)) => (e20 = inv2(e21))) & ((e20 = inv2(e20)) => (e20 = inv2(e20)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG027+1.p', ax15)).
fof(f4943, plain, (~ spl150_502 | spl150_495), inference(avatar_split_clause, [], [f1010, f4805, f4834])).
fof(f4834, plain, (spl150_502 <=> (e23 = inv2(e22))), introduced(avatar_definition, [new_symbols(naming, [spl150_502])])).
fof(f1010, plain, ((e22 = inv2(e23)) | ~ (e23 = inv2(e22))), inference(cnf_transformation, [], [f59])).
fof(f4941, plain, (~ spl150_504 | spl150_483), inference(avatar_split_clause, [], [f1012, f4755, f4842])).
fof(f4842, plain, (spl150_504 <=> (e25 = inv2(e22))), introduced(avatar_definition, [new_symbols(naming, [spl150_504])])).
fof(f1012, plain, ((e22 = inv2(e25)) | ~ (e25 = inv2(e22))), inference(cnf_transformation, [], [f59])).
fof(f4939, plain, (~ spl150_494 | spl150_508), inference(avatar_split_clause, [], [f1014, f4859, f4801])).
fof(f4801, plain, (spl150_494 <=> (e21 = inv2(e23))), introduced(avatar_definition, [new_symbols(naming, [spl150_494])])).
fof(f1014, plain, ((e23 = inv2(e21)) | ~ (e21 = inv2(e23))), inference(cnf_transformation, [], [f59])).
fof(f4927, plain, (~ spl150_484 | spl150_498), inference(avatar_split_clause, [], [f1028, f4817, f4759])).
fof(f4759, plain, (spl150_484 <=> (e23 = inv2(e25))), introduced(avatar_definition, [new_symbols(naming, [spl150_484])])).
fof(f1028, plain, ((e25 = inv2(e23)) | ~ (e23 = inv2(e25))), inference(cnf_transformation, [], [f59])).
fof(f4870, plain, (spl150_505 | spl150_506 | spl150_507 | spl150_508 | spl150_509 | spl150_510), inference(avatar_split_clause, [], [f925, f4867, f4863, f4859, f4855, f4851, f4847])).
fof(f925, plain, ((e25 = inv2(e21)) | (e24 = inv2(e21)) | (e23 = inv2(e21)) | (e22 = inv2(e21)) | (e21 = inv2(e21)) | (e20 = inv2(e21))), inference(cnf_transformation, [], [f8])).
fof(f4845, plain, (spl150_499 | spl150_500 | spl150_501 | spl150_502 | spl150_503 | spl150_504), inference(avatar_split_clause, [], [f926, f4842, f4838, f4834, f4830, f4826, f4822])).
fof(f926, plain, ((e25 = inv2(e22)) | (e24 = inv2(e22)) | (e23 = inv2(e22)) | (e22 = inv2(e22)) | (e21 = inv2(e22)) | (e20 = inv2(e22))), inference(cnf_transformation, [], [f8])).
fof(f4820, plain, (spl150_493 | spl150_494 | spl150_495 | spl150_496 | spl150_497 | spl150_498), inference(avatar_split_clause, [], [f927, f4817, f4813, f4809, f4805, f4801, f4797])).
fof(f927, plain, ((e25 = inv2(e23)) | (e24 = inv2(e23)) | (e23 = inv2(e23)) | (e22 = inv2(e23)) | (e21 = inv2(e23)) | (e20 = inv2(e23))), inference(cnf_transformation, [], [f8])).
fof(f4770, plain, (spl150_481 | spl150_482 | spl150_483 | spl150_484 | spl150_485 | spl150_486), inference(avatar_split_clause, [], [f929, f4767, f4763, f4759, f4755, f4751, f4747])).
fof(f929, plain, ((e25 = inv2(e25)) | (e24 = inv2(e25)) | (e23 = inv2(e25)) | (e22 = inv2(e25)) | (e21 = inv2(e25)) | (e20 = inv2(e25))), inference(cnf_transformation, [], [f8])).
fof(f4745, plain, (spl150_475 | spl150_476 | spl150_477 | spl150_478 | spl150_479 | spl150_480), inference(avatar_split_clause, [], [f911, f4742, f4738, f4734, f4730, f4726, f4722])).
fof(f911, plain, ((e25 = unit2) | (e24 = unit2) | (e23 = unit2) | (e22 = unit2) | (e21 = unit2) | (e20 = unit2)), inference(cnf_transformation, [], [f7])).
fof(f4470, plain, (spl150_409 | spl150_410 | spl150_411 | spl150_412 | spl150_413 | spl150_414), inference(avatar_split_clause, [], [f657, f4467, f4463, f4459, f4455, f4451, f4447])).
fof(f657, plain, ((e25 = op2(e21, e24)) | (e24 = op2(e21, e24)) | (e23 = op2(e21, e24)) | (e22 = op2(e21, e24)) | (e21 = op2(e21, e24)) | (e20 = op2(e21, e24))), inference(cnf_transformation, [], [f5])).
fof(f5, plain, (((e25 = op2(e25, e25)) | (e24 = op2(e25, e25)) | (e23 = op2(e25, e25)) | (e22 = op2(e25, e25)) | (e21 = op2(e25, e25)) | (e20 = op2(e25, e25))) & ((e25 = op2(e25, e24)) | (e24 = op2(e25, e24)) | (e23 = op2(e25, e24)) | (e22 = op2(e25, e24)) | (e21 = op2(e25, e24)) | (e20 = op2(e25, e24))) & ((e25 = op2(e25, e23)) | (e24 = op2(e25, e23)) | (e23 = op2(e25, e23)) | (e22 = op2(e25, e23)) | (e21 = op2(e25, e23)) | (e20 = op2(e25, e23))) & ((e25 = op2(e25, e22)) | (e24 = op2(e25, e22)) | (e23 = op2(e25, e22)) | (e22 = op2(e25, e22)) | (e21 = op2(e25, e22)) | (e20 = op2(e25, e22))) & ((e25 = op2(e25, e21)) | (e24 = op2(e25, e21)) | (e23 = op2(e25, e21)) | (e22 = op2(e25, e21)) | (e21 = op2(e25, e21)) | (e20 = op2(e25, e21))) & ((e25 = op2(e25, e20)) | (e24 = op2(e25, e20)) | (e23 = op2(e25, e20)) | (e22 = op2(e25, e20)) | (e21 = op2(e25, e20)) | (e20 = op2(e25, e20))) & ((e25 = op2(e24, e25)) | (e24 = op2(e24, e25)) | (e23 = op2(e24, e25)) | (e22 = op2(e24, e25)) | (e21 = op2(e24, e25)) | (e20 = op2(e24, e25))) & ((e25 = op2(e24, e24)) | (e24 = op2(e24, e24)) | (e23 = op2(e24, e24)) | (e22 = op2(e24, e24)) | (e21 = op2(e24, e24)) | (e20 = op2(e24, e24))) & ((e25 = op2(e24, e23)) | (e24 = op2(e24, e23)) | (e23 = op2(e24, e23)) | (e22 = op2(e24, e23)) | (e21 = op2(e24, e23)) | (e20 = op2(e24, e23))) & ((e25 = op2(e24, e22)) | (e24 = op2(e24, e22)) | (e23 = op2(e24, e22)) | (e22 = op2(e24, e22)) | (e21 = op2(e24, e22)) | (e20 = op2(e24, e22))) & ((e25 = op2(e24, e21)) | (e24 = op2(e24, e21)) | (e23 = op2(e24, e21)) | (e22 = op2(e24, e21)) | (e21 = op2(e24, e21)) | (e20 = op2(e24, e21))) & ((e25 = op2(e24, e20)) | (e24 = op2(e24, e20)) | (e23 = op2(e24, e20)) | (e22 = op2(e24, e20)) | (e21 = op2(e24, e20)) | (e20 = op2(e24, e20))) & ((e25 = op2(e23, e25)) | (e24 = op2(e23, e25)) | (e23 = op2(e23, e25)) | (e22 = op2(e23, e25)) | (e21 = op2(e23, e25)) | (e20 = op2(e23, e25))) & ((e25 = op2(e23, e24)) | (e24 = op2(e23, e24)) | (e23 = op2(e23, e24)) | (e22 = op2(e23, e24)) | (e21 = op2(e23, e24)) | (e20 = op2(e23, e24))) & ((e25 = op2(e23, e23)) | (e24 = op2(e23, e23)) | (e23 = op2(e23, e23)) | (e22 = op2(e23, e23)) | (e21 = op2(e23, e23)) | (e20 = op2(e23, e23))) & ((e25 = op2(e23, e22)) | (e24 = op2(e23, e22)) | (e23 = op2(e23, e22)) | (e22 = op2(e23, e22)) | (e21 = op2(e23, e22)) | (e20 = op2(e23, e22))) & ((e25 = op2(e23, e21)) | (e24 = op2(e23, e21)) | (e23 = op2(e23, e21)) | (e22 = op2(e23, e21)) | (e21 = op2(e23, e21)) | (e20 = op2(e23, e21))) & ((e25 = op2(e23, e20)) | (e24 = op2(e23, e20)) | (e23 = op2(e23, e20)) | (e22 = op2(e23, e20)) | (e21 = op2(e23, e20)) | (e20 = op2(e23, e20))) & ((e25 = op2(e22, e25)) | (e24 = op2(e22, e25)) | (e23 = op2(e22, e25)) | (e22 = op2(e22, e25)) | (e21 = op2(e22, e25)) | (e20 = op2(e22, e25))) & ((e25 = op2(e22, e24)) | (e24 = op2(e22, e24)) | (e23 = op2(e22, e24)) | (e22 = op2(e22, e24)) | (e21 = op2(e22, e24)) | (e20 = op2(e22, e24))) & ((e25 = op2(e22, e23)) | (e24 = op2(e22, e23)) | (e23 = op2(e22, e23)) | (e22 = op2(e22, e23)) | (e21 = op2(e22, e23)) | (e20 = op2(e22, e23))) & ((e25 = op2(e22, e22)) | (e24 = op2(e22, e22)) | (e23 = op2(e22, e22)) | (e22 = op2(e22, e22)) | (e21 = op2(e22, e22)) | (e20 = op2(e22, e22))) & ((e25 = op2(e22, e21)) | (e24 = op2(e22, e21)) | (e23 = op2(e22, e21)) | (e22 = op2(e22, e21)) | (e21 = op2(e22, e21)) | (e20 = op2(e22, e21))) & ((e25 = op2(e22, e20)) | (e24 = op2(e22, e20)) | (e23 = op2(e22, e20)) | (e22 = op2(e22, e20)) | (e21 = op2(e22, e20)) | (e20 = op2(e22, e20))) & ((e25 = op2(e21, e25)) | (e24 = op2(e21, e25)) | (e23 = op2(e21, e25)) | (e22 = op2(e21, e25)) | (e21 = op2(e21, e25)) | (e20 = op2(e21, e25))) & ((e25 = op2(e21, e24)) | (e24 = op2(e21, e24)) | (e23 = op2(e21, e24)) | (e22 = op2(e21, e24)) | (e21 = op2(e21, e24)) | (e20 = op2(e21, e24))) & ((e25 = op2(e21, e23)) | (e24 = op2(e21, e23)) | (e23 = op2(e21, e23)) | (e22 = op2(e21, e23)) | (e21 = op2(e21, e23)) | (e20 = op2(e21, e23))) & ((e25 = op2(e21, e22)) | (e24 = op2(e21, e22)) | (e23 = op2(e21, e22)) | (e22 = op2(e21, e22)) | (e21 = op2(e21, e22)) | (e20 = op2(e21, e22))) & ((e25 = op2(e21, e21)) | (e24 = op2(e21, e21)) | (e23 = op2(e21, e21)) | (e22 = op2(e21, e21)) | (e21 = op2(e21, e21)) | (e20 = op2(e21, e21))) & ((e25 = op2(e21, e20)) | (e24 = op2(e21, e20)) | (e23 = op2(e21, e20)) | (e22 = op2(e21, e20)) | (e21 = op2(e21, e20)) | (e20 = op2(e21, e20))) & ((e25 = op2(e20, e25)) | (e24 = op2(e20, e25)) | (e23 = op2(e20, e25)) | (e22 = op2(e20, e25)) | (e21 = op2(e20, e25)) | (e20 = op2(e20, e25))) & ((e25 = op2(e20, e24)) | (e24 = op2(e20, e24)) | (e23 = op2(e20, e24)) | (e22 = op2(e20, e24)) | (e21 = op2(e20, e24)) | (e20 = op2(e20, e24))) & ((e25 = op2(e20, e23)) | (e24 = op2(e20, e23)) | (e23 = op2(e20, e23)) | (e22 = op2(e20, e23)) | (e21 = op2(e20, e23)) | (e20 = op2(e20, e23))) & ((e25 = op2(e20, e22)) | (e24 = op2(e20, e22)) | (e23 = op2(e20, e22)) | (e22 = op2(e20, e22)) | (e21 = op2(e20, e22)) | (e20 = op2(e20, e22))) & ((e25 = op2(e20, e21)) | (e24 = op2(e20, e21)) | (e23 = op2(e20, e21)) | (e22 = op2(e20, e21)) | (e21 = op2(e20, e21)) | (e20 = op2(e20, e21))) & ((op2(e20, e20) = e25) | (op2(e20, e20) = e24) | (op2(e20, e20) = e23) | (op2(e20, e20) = e22) | (op2(e20, e20) = e21) | (e20 = op2(e20, e20)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG027+1.p', ax5)).
fof(f4445, plain, (spl150_403 | spl150_404 | spl150_405 | spl150_406 | spl150_407 | spl150_408), inference(avatar_split_clause, [], [f658, f4442, f4438, f4434, f4430, f4426, f4422])).
fof(f658, plain, ((e25 = op2(e21, e25)) | (e24 = op2(e21, e25)) | (e23 = op2(e21, e25)) | (e22 = op2(e21, e25)) | (e21 = op2(e21, e25)) | (e20 = op2(e21, e25))), inference(cnf_transformation, [], [f5])).
fof(f4095, plain, (spl150_319 | spl150_320 | spl150_321 | spl150_322 | spl150_323 | spl150_324), inference(avatar_split_clause, [], [f672, f4092, f4088, f4084, f4080, f4076, f4072])).
fof(f672, plain, ((e25 = op2(e24, e21)) | (e24 = op2(e24, e21)) | (e23 = op2(e24, e21)) | (e22 = op2(e24, e21)) | (e21 = op2(e24, e21)) | (e20 = op2(e24, e21))), inference(cnf_transformation, [], [f5])).
fof(f4070, plain, (spl150_313 | spl150_314 | spl150_315 | spl150_316 | spl150_317 | spl150_318), inference(avatar_split_clause, [], [f673, f4067, f4063, f4059, f4055, f4051, f4047])).
fof(f673, plain, ((e25 = op2(e24, e22)) | (e24 = op2(e24, e22)) | (e23 = op2(e24, e22)) | (e22 = op2(e24, e22)) | (e21 = op2(e24, e22)) | (e20 = op2(e24, e22))), inference(cnf_transformation, [], [f5])).
fof(f3995, plain, (spl150_295 | spl150_296 | spl150_297 | spl150_298 | spl150_299 | spl150_300), inference(avatar_split_clause, [], [f676, f3992, f3988, f3984, f3980, f3976, f3972])).
fof(f676, plain, ((e25 = op2(e24, e25)) | (e24 = op2(e24, e25)) | (e23 = op2(e24, e25)) | (e22 = op2(e24, e25)) | (e21 = op2(e24, e25)) | (e20 = op2(e24, e25))), inference(cnf_transformation, [], [f5])).
fof(f3795, plain, (spl150_247 | spl150_248 | spl150_249 | spl150_250 | spl150_251 | spl150_252), inference(avatar_split_clause, [], [f642, f3792, f3788, f3784, f3780, f3776, f3772])).
fof(f642, plain, ((e15 = inv1(e11)) | (e14 = inv1(e11)) | (e13 = inv1(e11)) | (e12 = inv1(e11)) | (e11 = inv1(e11)) | (e10 = inv1(e11))), inference(cnf_transformation, [], [f4])).
fof(f3670, plain, (spl150_217 | spl150_218 | spl150_219 | spl150_220 | spl150_221 | spl150_222), inference(avatar_split_clause, [], [f628, f3667, f3663, f3659, f3655, f3651, f3647])).
fof(f628, plain, ((e15 = unit1) | (e14 = unit1) | (e13 = unit1) | (e12 = unit1) | (e11 = unit1) | (e10 = unit1)), inference(cnf_transformation, [], [f3])).
fof(f3420, plain, (spl150_157 | spl150_158 | spl150_159 | spl150_160 | spl150_161 | spl150_162), inference(avatar_split_clause, [], [f373, f3417, f3413, f3409, f3405, f3401, f3397])).
fof(f373, plain, ((e15 = op1(e11, e13)) | (e14 = op1(e11, e13)) | (e13 = op1(e11, e13)) | (e12 = op1(e11, e13)) | (e11 = op1(e11, e13)) | (e10 = op1(e11, e13))), inference(cnf_transformation, [], [f1])).
fof(f1, plain, (((e15 = op1(e15, e15)) | (e14 = op1(e15, e15)) | (e13 = op1(e15, e15)) | (e12 = op1(e15, e15)) | (e11 = op1(e15, e15)) | (e10 = op1(e15, e15))) & ((e15 = op1(e15, e14)) | (e14 = op1(e15, e14)) | (e13 = op1(e15, e14)) | (e12 = op1(e15, e14)) | (e11 = op1(e15, e14)) | (e10 = op1(e15, e14))) & ((e15 = op1(e15, e13)) | (e14 = op1(e15, e13)) | (e13 = op1(e15, e13)) | (e12 = op1(e15, e13)) | (e11 = op1(e15, e13)) | (e10 = op1(e15, e13))) & ((e15 = op1(e15, e12)) | (e14 = op1(e15, e12)) | (e13 = op1(e15, e12)) | (e12 = op1(e15, e12)) | (e11 = op1(e15, e12)) | (e10 = op1(e15, e12))) & ((e15 = op1(e15, e11)) | (e14 = op1(e15, e11)) | (e13 = op1(e15, e11)) | (e12 = op1(e15, e11)) | (e11 = op1(e15, e11)) | (e10 = op1(e15, e11))) & ((e15 = op1(e15, e10)) | (e14 = op1(e15, e10)) | (e13 = op1(e15, e10)) | (e12 = op1(e15, e10)) | (e11 = op1(e15, e10)) | (e10 = op1(e15, e10))) & ((e15 = op1(e14, e15)) | (e14 = op1(e14, e15)) | (e13 = op1(e14, e15)) | (e12 = op1(e14, e15)) | (e11 = op1(e14, e15)) | (e10 = op1(e14, e15))) & ((e15 = op1(e14, e14)) | (e14 = op1(e14, e14)) | (e13 = op1(e14, e14)) | (e12 = op1(e14, e14)) | (e11 = op1(e14, e14)) | (e10 = op1(e14, e14))) & ((e15 = op1(e14, e13)) | (e14 = op1(e14, e13)) | (e13 = op1(e14, e13)) | (e12 = op1(e14, e13)) | (e11 = op1(e14, e13)) | (e10 = op1(e14, e13))) & ((e15 = op1(e14, e12)) | (e14 = op1(e14, e12)) | (e13 = op1(e14, e12)) | (e12 = op1(e14, e12)) | (e11 = op1(e14, e12)) | (e10 = op1(e14, e12))) & ((e15 = op1(e14, e11)) | (e14 = op1(e14, e11)) | (e13 = op1(e14, e11)) | (e12 = op1(e14, e11)) | (e11 = op1(e14, e11)) | (e10 = op1(e14, e11))) & ((e15 = op1(e14, e10)) | (e14 = op1(e14, e10)) | (e13 = op1(e14, e10)) | (e12 = op1(e14, e10)) | (e11 = op1(e14, e10)) | (e10 = op1(e14, e10))) & ((e15 = op1(e13, e15)) | (e14 = op1(e13, e15)) | (e13 = op1(e13, e15)) | (e12 = op1(e13, e15)) | (e11 = op1(e13, e15)) | (e10 = op1(e13, e15))) & ((e15 = op1(e13, e14)) | (e14 = op1(e13, e14)) | (e13 = op1(e13, e14)) | (e12 = op1(e13, e14)) | (e11 = op1(e13, e14)) | (e10 = op1(e13, e14))) & ((e15 = op1(e13, e13)) | (e14 = op1(e13, e13)) | (e13 = op1(e13, e13)) | (e12 = op1(e13, e13)) | (e11 = op1(e13, e13)) | (e10 = op1(e13, e13))) & ((e15 = op1(e13, e12)) | (e14 = op1(e13, e12)) | (e13 = op1(e13, e12)) | (e12 = op1(e13, e12)) | (e11 = op1(e13, e12)) | (e10 = op1(e13, e12))) & ((e15 = op1(e13, e11)) | (e14 = op1(e13, e11)) | (e13 = op1(e13, e11)) | (e12 = op1(e13, e11)) | (e11 = op1(e13, e11)) | (e10 = op1(e13, e11))) & ((e15 = op1(e13, e10)) | (e14 = op1(e13, e10)) | (e13 = op1(e13, e10)) | (e12 = op1(e13, e10)) | (e11 = op1(e13, e10)) | (e10 = op1(e13, e10))) & ((e15 = op1(e12, e15)) | (e14 = op1(e12, e15)) | (e13 = op1(e12, e15)) | (e12 = op1(e12, e15)) | (e11 = op1(e12, e15)) | (e10 = op1(e12, e15))) & ((e15 = op1(e12, e14)) | (e14 = op1(e12, e14)) | (e13 = op1(e12, e14)) | (e12 = op1(e12, e14)) | (e11 = op1(e12, e14)) | (e10 = op1(e12, e14))) & ((e15 = op1(e12, e13)) | (e14 = op1(e12, e13)) | (e13 = op1(e12, e13)) | (e12 = op1(e12, e13)) | (e11 = op1(e12, e13)) | (e10 = op1(e12, e13))) & ((e15 = op1(e12, e12)) | (e14 = op1(e12, e12)) | (e13 = op1(e12, e12)) | (e12 = op1(e12, e12)) | (e11 = op1(e12, e12)) | (e10 = op1(e12, e12))) & ((e15 = op1(e12, e11)) | (e14 = op1(e12, e11)) | (e13 = op1(e12, e11)) | (e12 = op1(e12, e11)) | (e11 = op1(e12, e11)) | (e10 = op1(e12, e11))) & ((e15 = op1(e12, e10)) | (e14 = op1(e12, e10)) | (e13 = op1(e12, e10)) | (e12 = op1(e12, e10)) | (e11 = op1(e12, e10)) | (e10 = op1(e12, e10))) & ((e15 = op1(e11, e15)) | (e14 = op1(e11, e15)) | (e13 = op1(e11, e15)) | (e12 = op1(e11, e15)) | (e11 = op1(e11, e15)) | (e10 = op1(e11, e15))) & ((e15 = op1(e11, e14)) | (e14 = op1(e11, e14)) | (e13 = op1(e11, e14)) | (e12 = op1(e11, e14)) | (e11 = op1(e11, e14)) | (e10 = op1(e11, e14))) & ((e15 = op1(e11, e13)) | (e14 = op1(e11, e13)) | (e13 = op1(e11, e13)) | (e12 = op1(e11, e13)) | (e11 = op1(e11, e13)) | (e10 = op1(e11, e13))) & ((e15 = op1(e11, e12)) | (e14 = op1(e11, e12)) | (e13 = op1(e11, e12)) | (e12 = op1(e11, e12)) | (e11 = op1(e11, e12)) | (e10 = op1(e11, e12))) & ((e15 = op1(e11, e11)) | (e14 = op1(e11, e11)) | (e13 = op1(e11, e11)) | (e12 = op1(e11, e11)) | (e11 = op1(e11, e11)) | (e10 = op1(e11, e11))) & ((e15 = op1(e11, e10)) | (e14 = op1(e11, e10)) | (e13 = op1(e11, e10)) | (e12 = op1(e11, e10)) | (e11 = op1(e11, e10)) | (e10 = op1(e11, e10))) & ((e15 = op1(e10, e15)) | (e14 = op1(e10, e15)) | (e13 = op1(e10, e15)) | (e12 = op1(e10, e15)) | (e11 = op1(e10, e15)) | (e10 = op1(e10, e15))) & ((e15 = op1(e10, e14)) | (e14 = op1(e10, e14)) | (e13 = op1(e10, e14)) | (e12 = op1(e10, e14)) | (e11 = op1(e10, e14)) | (e10 = op1(e10, e14))) & ((e15 = op1(e10, e13)) | (e14 = op1(e10, e13)) | (e13 = op1(e10, e13)) | (e12 = op1(e10, e13)) | (e11 = op1(e10, e13)) | (e10 = op1(e10, e13))) & ((e15 = op1(e10, e12)) | (e14 = op1(e10, e12)) | (e13 = op1(e10, e12)) | (e12 = op1(e10, e12)) | (e11 = op1(e10, e12)) | (e10 = op1(e10, e12))) & ((e15 = op1(e10, e11)) | (e14 = op1(e10, e11)) | (e13 = op1(e10, e11)) | (e12 = op1(e10, e11)) | (e11 = op1(e10, e11)) | (e10 = op1(e10, e11))) & ((op1(e10, e10) = e15) | (op1(e10, e10) = e14) | (op1(e10, e10) = e13) | (op1(e10, e10) = e12) | (op1(e10, e10) = e11) | (e10 = op1(e10, e10)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG027+1.p', ax1)).
fof(f3395, plain, (spl150_151 | spl150_152 | spl150_153 | spl150_154 | spl150_155 | spl150_156), inference(avatar_split_clause, [], [f374, f3392, f3388, f3384, f3380, f3376, f3372])).
fof(f374, plain, ((e15 = op1(e11, e14)) | (e14 = op1(e11, e14)) | (e13 = op1(e11, e14)) | (e12 = op1(e11, e14)) | (e11 = op1(e11, e14)) | (e10 = op1(e11, e14))), inference(cnf_transformation, [], [f1])).
fof(f3320, plain, (spl150_133 | spl150_134 | spl150_135 | spl150_136 | spl150_137 | spl150_138), inference(avatar_split_clause, [], [f377, f3317, f3313, f3309, f3305, f3301, f3297])).
fof(f377, plain, ((e15 = op1(e12, e11)) | (e14 = op1(e12, e11)) | (e13 = op1(e12, e11)) | (e12 = op1(e12, e11)) | (e11 = op1(e12, e11)) | (e10 = op1(e12, e11))), inference(cnf_transformation, [], [f1])).
fof(f3270, plain, (spl150_121 | spl150_122 | spl150_123 | spl150_124 | spl150_125 | spl150_126), inference(avatar_split_clause, [], [f379, f3267, f3263, f3259, f3255, f3251, f3247])).
fof(f379, plain, ((e15 = op1(e12, e13)) | (e14 = op1(e12, e13)) | (e13 = op1(e12, e13)) | (e12 = op1(e12, e13)) | (e11 = op1(e12, e13)) | (e10 = op1(e12, e13))), inference(cnf_transformation, [], [f1])).
fof(f2920, plain, (spl150_37 | spl150_38 | spl150_39 | spl150_40 | spl150_41 | spl150_42), inference(avatar_split_clause, [], [f393, f2917, f2913, f2909, f2905, f2901, f2897])).
fof(f393, plain, ((e15 = op1(e14, e15)) | (e14 = op1(e14, e15)) | (e13 = op1(e14, e15)) | (e12 = op1(e14, e15)) | (e11 = op1(e14, e15)) | (e10 = op1(e14, e15))), inference(cnf_transformation, [], [f1])).
fof(f2820, plain, (spl150_13 | spl150_14 | spl150_15 | spl150_16 | spl150_17 | spl150_18), inference(avatar_split_clause, [], [f397, f2817, f2813, f2809, f2805, f2801, f2797])).
fof(f397, plain, ((e15 = op1(e15, e13)) | (e14 = op1(e15, e13)) | (e13 = op1(e15, e13)) | (e12 = op1(e15, e13)) | (e11 = op1(e15, e13)) | (e10 = op1(e15, e13))), inference(cnf_transformation, [], [f1])).