fof(f14430, plain, $false, inference(avatar_sat_refutation, [], [f1356, f1818, f1823, f1825, f1827, f1829, f1833, f1834, f1838, f1839, f1846, f1848, f1849, f1852, f2414, f2421, f2423, f2429, f2431, f2433, f2435, f2436, f2444, f2446, f2447, f2448, f2449, f2450, f2466, f2469, f2470, f2473, f2475, f2477, f2478, f2480, f2490, f2491, f2493, f2497, f2498, f2499, f2500, f2505, f2506, f2507, f2508, f2801, f2831, f2860, f4641, f8383, f8384, f8385, f8401, f8402, f8404, f8415, f8422, f8424, f8466, f8477, f8488, f8493, f8538, f8541, f8635, f8689, f8701, f8707, f8708, f8737, f8738, f8743, f8809, f8907, f8945, f8948, f8953, f8954, f8955, f9046, f9118, f9136, f9142, f9152, f9181, f9183, f9199, f9238, f9255, f9260, f9272, f9301, f9389, f9420, f9426, f9428, f9478, f9511, f9517, f9519, f9524, f9539, f9581, f9602, f9679, f9681, f9684, f9691, f9709, f9731, f9747, f9794, f9813, f9814, f9817, f9835, f9836, f9853, f9869, f9872, f9875, f9896, f9923, f9928, f9949, f10016, f10019, f10030, f10058, f10081, f10096, f10097, f10100, f10115, f10121, f10202, f10242, f10336, f10352, f10369, f10373, f10376, f10381, f10430, f10431, f10524, f10540, f10548, f10555, f10629, f10641, f10654, f10663, f10693, f10699, f10707, f10778, f10962, f11007, f11050, f11051, f11052, f11053, f11056, f11059, f11060, f11186, f11193, f11199, f11313, f11386, f11387, f11388, f11389, f11390, f11429, f11467, f11468, f11493, f11522, f11549, f11638, f11648, f11780, f11822, f11828, f11831, f11974, f12142, f12272, f12403, f12532, f12660, f12769, f12867, f12956, f13040, f13119, f13214, f13306, f13397, f13486, f13568, f13659, f13748, f13832, f13926, f14018, f14112, f14143, f14219, f14240, f14304, f14374, f14429])).
fof(f14429, plain, (~ spl80_27 | ~ spl80_29), inference(avatar_contradiction_clause, [], [f14428])).
fof(f14428, plain, ($false | (~ spl80_27 | ~ spl80_29)), inference(subsumption_resolution, [], [f14427, f577])).
fof(f577, plain, ~ (e11 = e13), inference(cnf_transformation, [], [f9])).
fof(f9, plain, (~ (e13 = e14) & ~ (e12 = e14) & ~ (e12 = e13) & ~ (e11 = e14) & ~ (e11 = e13) & ~ (e11 = e12) & ~ (e10 = e14) & ~ (e10 = e13) & ~ (e10 = e12) & ~ (e10 = e11)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG049+1.p', ax9)).
fof(f14427, plain, ((e11 = e13) | (~ spl80_27 | ~ spl80_29)), inference(forward_demodulation, [], [f1393, f1385])).
fof(f1385, plain, ((e11 = op1(e13, e14)) | ~ spl80_27), inference(avatar_component_clause, [], [f1383])).
fof(f1383, plain, (spl80_27 <=> (e11 = op1(e13, e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_27])])).
fof(f1393, plain, ((e13 = op1(e13, e14)) | ~ spl80_29), inference(avatar_component_clause, [], [f1391])).
fof(f1391, plain, (spl80_29 <=> (e13 = op1(e13, e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_29])])).
fof(f14374, plain, (~ spl80_113 | ~ spl80_115), inference(avatar_contradiction_clause, [], [f14373])).
fof(f14373, plain, ($false | (~ spl80_113 | ~ spl80_115)), inference(subsumption_resolution, [], [f14372, f580])).
fof(f580, plain, ~ (e12 = e14), inference(cnf_transformation, [], [f9])).
fof(f14372, plain, ((e12 = e14) | (~ spl80_113 | ~ spl80_115)), inference(forward_demodulation, [], [f1754, f1746])).
fof(f1746, plain, ((e12 = op1(e10, e12)) | ~ spl80_113), inference(avatar_component_clause, [], [f1744])).
fof(f1744, plain, (spl80_113 <=> (e12 = op1(e10, e12))), introduced(avatar_definition, [new_symbols(naming, [spl80_113])])).
fof(f1754, plain, ((e14 = op1(e10, e12)) | ~ spl80_115), inference(avatar_component_clause, [], [f1752])).
fof(f1752, plain, (spl80_115 <=> (e14 = op1(e10, e12))), introduced(avatar_definition, [new_symbols(naming, [spl80_115])])).
fof(f14304, plain, (~ spl80_39 | ~ spl80_49), inference(avatar_split_clause, [], [f14303, f1475, f1433])).
fof(f1433, plain, (spl80_39 <=> (e13 = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl80_39])])).
fof(f1475, plain, (spl80_49 <=> (e13 = op1(e13, e10))), introduced(avatar_definition, [new_symbols(naming, [spl80_49])])).
fof(f14303, plain, (~ (e13 = op1(e13, e12)) | ~ spl80_49), inference(forward_demodulation, [], [f453, f1477])).
fof(f1477, plain, ((e13 = op1(e13, e10)) | ~ spl80_49), inference(avatar_component_clause, [], [f1475])).
fof(f453, plain, ~ (op1(e13, e10) = op1(e13, e12)), inference(cnf_transformation, [], [f7])).
fof(f7, plain, (~ (op1(e14, e13) = op1(e14, e14)) & ~ (op1(e14, e12) = op1(e14, e14)) & ~ (op1(e14, e12) = op1(e14, e13)) & ~ (op1(e14, e11) = op1(e14, e14)) & ~ (op1(e14, e11) = op1(e14, e13)) & ~ (op1(e14, e11) = op1(e14, e12)) & ~ (op1(e14, e10) = op1(e14, e14)) & ~ (op1(e14, e10) = op1(e14, e13)) & ~ (op1(e14, e10) = op1(e14, e12)) & ~ (op1(e14, e10) = op1(e14, e11)) & ~ (op1(e13, e13) = op1(e13, e14)) & ~ (op1(e13, e12) = op1(e13, e14)) & ~ (op1(e13, e12) = op1(e13, e13)) & ~ (op1(e13, e11) = op1(e13, e14)) & ~ (op1(e13, e11) = op1(e13, e13)) & ~ (op1(e13, e11) = op1(e13, e12)) & ~ (op1(e13, e10) = op1(e13, e14)) & ~ (op1(e13, e10) = op1(e13, e13)) & ~ (op1(e13, e10) = op1(e13, e12)) & ~ (op1(e13, e10) = op1(e13, e11)) & ~ (op1(e12, e13) = op1(e12, e14)) & ~ (op1(e12, e12) = op1(e12, e14)) & ~ (op1(e12, e12) = op1(e12, e13)) & ~ (op1(e12, e11) = op1(e12, e14)) & ~ (op1(e12, e11) = op1(e12, e13)) & ~ (op1(e12, e11) = op1(e12, e12)) & ~ (op1(e12, e10) = op1(e12, e14)) & ~ (op1(e12, e10) = op1(e12, e13)) & ~ (op1(e12, e10) = op1(e12, e12)) & ~ (op1(e12, e10) = op1(e12, e11)) & ~ (op1(e11, e13) = op1(e11, e14)) & ~ (op1(e11, e12) = op1(e11, e14)) & ~ (op1(e11, e12) = op1(e11, e13)) & ~ (op1(e11, e11) = op1(e11, e14)) & ~ (op1(e11, e11) = op1(e11, e13)) & ~ (op1(e11, e11) = op1(e11, e12)) & ~ (op1(e11, e10) = op1(e11, e14)) & ~ (op1(e11, e10) = op1(e11, e13)) & ~ (op1(e11, e10) = op1(e11, e12)) & ~ (op1(e11, e10) = op1(e11, e11)) & ~ (op1(e10, e13) = op1(e10, e14)) & ~ (op1(e10, e12) = op1(e10, e14)) & ~ (op1(e10, e12) = op1(e10, e13)) & ~ (op1(e10, e11) = op1(e10, e14)) & ~ (op1(e10, e11) = op1(e10, e13)) & ~ (op1(e10, e11) = op1(e10, e12)) & ~ (op1(e10, e10) = op1(e10, e14)) & ~ (op1(e10, e10) = op1(e10, e13)) & ~ (op1(e10, e10) = op1(e10, e12)) & ~ (op1(e10, e10) = op1(e10, e11)) & ~ (op1(e13, e14) = op1(e14, e14)) & ~ (op1(e12, e14) = op1(e14, e14)) & ~ (op1(e12, e14) = op1(e13, e14)) & ~ (op1(e11, e14) = op1(e14, e14)) & ~ (op1(e11, e14) = op1(e13, e14)) & ~ (op1(e11, e14) = op1(e12, e14)) & ~ (op1(e10, e14) = op1(e14, e14)) & ~ (op1(e10, e14) = op1(e13, e14)) & ~ (op1(e10, e14) = op1(e12, e14)) & ~ (op1(e10, e14) = op1(e11, e14)) & ~ (op1(e13, e13) = op1(e14, e13)) & ~ (op1(e12, e13) = op1(e14, e13)) & ~ (op1(e12, e13) = op1(e13, e13)) & ~ (op1(e11, e13) = op1(e14, e13)) & ~ (op1(e11, e13) = op1(e13, e13)) & ~ (op1(e11, e13) = op1(e12, e13)) & ~ (op1(e10, e13) = op1(e14, e13)) & ~ (op1(e10, e13) = op1(e13, e13)) & ~ (op1(e10, e13) = op1(e12, e13)) & ~ (op1(e10, e13) = op1(e11, e13)) & ~ (op1(e13, e12) = op1(e14, e12)) & ~ (op1(e12, e12) = op1(e14, e12)) & ~ (op1(e12, e12) = op1(e13, e12)) & ~ (op1(e11, e12) = op1(e14, e12)) & ~ (op1(e11, e12) = op1(e13, e12)) & ~ (op1(e11, e12) = op1(e12, e12)) & ~ (op1(e10, e12) = op1(e14, e12)) & ~ (op1(e10, e12) = op1(e13, e12)) & ~ (op1(e10, e12) = op1(e12, e12)) & ~ (op1(e10, e12) = op1(e11, e12)) & ~ (op1(e13, e11) = op1(e14, e11)) & ~ (op1(e12, e11) = op1(e14, e11)) & ~ (op1(e12, e11) = op1(e13, e11)) & ~ (op1(e11, e11) = op1(e14, e11)) & ~ (op1(e11, e11) = op1(e13, e11)) & ~ (op1(e11, e11) = op1(e12, e11)) & ~ (op1(e10, e11) = op1(e14, e11)) & ~ (op1(e10, e11) = op1(e13, e11)) & ~ (op1(e10, e11) = op1(e12, e11)) & ~ (op1(e10, e11) = op1(e11, e11)) & ~ (op1(e13, e10) = op1(e14, e10)) & ~ (op1(e12, e10) = op1(e14, e10)) & ~ (op1(e12, e10) = op1(e13, e10)) & ~ (op1(e11, e10) = op1(e14, e10)) & ~ (op1(e11, e10) = op1(e13, e10)) & ~ (op1(e11, e10) = op1(e12, e10)) & ~ (op1(e10, e10) = op1(e14, e10)) & ~ (op1(e10, e10) = op1(e13, e10)) & ~ (op1(e10, e10) = op1(e12, e10)) & ~ (op1(e10, e10) = op1(e11, e10))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG049+1.p', ax7)).
fof(f14240, plain, (~ spl80_117 | ~ spl80_120), inference(avatar_contradiction_clause, [], [f14239])).
fof(f14239, plain, ($false | (~ spl80_117 | ~ spl80_120)), inference(subsumption_resolution, [], [f14238, f578])).
fof(f578, plain, ~ (e11 = e14), inference(cnf_transformation, [], [f9])).
fof(f14238, plain, ((e11 = e14) | (~ spl80_117 | ~ spl80_120)), inference(forward_demodulation, [], [f1775, f1763])).
fof(f1763, plain, ((e11 = op1(e10, e11)) | ~ spl80_117), inference(avatar_component_clause, [], [f1761])).
fof(f1761, plain, (spl80_117 <=> (e11 = op1(e10, e11))), introduced(avatar_definition, [new_symbols(naming, [spl80_117])])).
fof(f1775, plain, ((e14 = op1(e10, e11)) | ~ spl80_120), inference(avatar_component_clause, [], [f1773])).
fof(f1773, plain, (spl80_120 <=> (e14 = op1(e10, e11))), introduced(avatar_definition, [new_symbols(naming, [spl80_120])])).
fof(f14219, plain, (~ spl80_97 | ~ spl80_203 | ~ spl80_291 | ~ spl80_296 | spl80_542), inference(avatar_contradiction_clause, [], [f14218])).
fof(f14218, plain, ($false | (~ spl80_97 | ~ spl80_203 | ~ spl80_291 | ~ spl80_296 | spl80_542)), inference(subsumption_resolution, [], [f14217, f2174])).
fof(f2174, plain, ((e22 = op2(e22, e20)) | ~ spl80_203), inference(avatar_component_clause, [], [f2172])).
fof(f2172, plain, (spl80_203 <=> (e22 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl80_203])])).
fof(f14217, plain, (~ (e22 = op2(e22, e20)) | (~ spl80_97 | ~ spl80_291 | ~ spl80_296 | spl80_542)), inference(forward_demodulation, [], [f14216, f11544])).
fof(f11544, plain, ((e22 = h18(e11)) | ~ spl80_291), inference(backward_demodulation, [], [f2643, f2829])).
fof(f2829, plain, ((e22 = h8(e13)) | ~ spl80_291), inference(avatar_component_clause, [], [f2828])).
fof(f2828, plain, (spl80_291 <=> (e22 = h8(e13))), introduced(avatar_definition, [new_symbols(naming, [spl80_291])])).
fof(f2643, plain, (h8(e13) = h18(e11)), inference(forward_demodulation, [], [f761, f712])).
fof(f712, plain, (op2(e21, e24) = h8(e13)), inference(cnf_transformation, [], [f23])).
fof(f23, plain, ((op2(e21, e24) = h8(e13)) & (op2(e24, e21) = h8(e11)) & (h8(e10) = op2(op2(e24, e21), op2(e24, e21))) & (e24 = h8(e14)) & (e21 = h8(e12))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG049+1.p', ax23)).
fof(f761, plain, (op2(e21, e24) = h18(e11)), inference(cnf_transformation, [], [f33])).
fof(f33, plain, ((op2(e24, e21) = h18(e13)) & (op2(e21, e24) = h18(e11)) & (h18(e10) = op2(op2(e21, e24), op2(e21, e24))) & (e21 = h18(e14)) & (e24 = h18(e12))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG049+1.p', ax33)).
fof(f14216, plain, (~ (op2(e22, e20) = h18(e11)) | (~ spl80_97 | ~ spl80_291 | ~ spl80_296 | spl80_542)), inference(forward_demodulation, [], [f14215, f1679])).
fof(f1679, plain, ((e11 = op1(e11, e10)) | ~ spl80_97), inference(avatar_component_clause, [], [f1677])).
fof(f1677, plain, (spl80_97 <=> (e11 = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl80_97])])).
fof(f14215, plain, (~ (op2(e22, e20) = h18(op1(e11, e10))) | (~ spl80_291 | ~ spl80_296 | spl80_542)), inference(forward_demodulation, [], [f14214, f2829])).
fof(f14214, plain, (~ (h18(op1(e11, e10)) = op2(h8(e13), e20)) | (~ spl80_296 | spl80_542)), inference(forward_demodulation, [], [f4581, f2858])).
fof(f2858, plain, ((e20 = h18(e10)) | ~ spl80_296), inference(avatar_component_clause, [], [f2857])).
fof(f2857, plain, (spl80_296 <=> (e20 = h18(e10))), introduced(avatar_definition, [new_symbols(naming, [spl80_296])])).
fof(f4581, plain, (~ (h18(op1(e11, e10)) = op2(h8(e13), h18(e10))) | spl80_542), inference(avatar_component_clause, [], [f4579])).
fof(f4579, plain, (spl80_542 <=> (h18(op1(e11, e10)) = op2(h8(e13), h18(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl80_542])])).
fof(f14143, plain, (~ spl80_67 | ~ spl80_117), inference(avatar_split_clause, [], [f14142, f1761, f1551])).
fof(f1551, plain, (spl80_67 <=> (e11 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl80_67])])).
fof(f14142, plain, (~ (e11 = op1(e12, e11)) | ~ spl80_117), inference(forward_demodulation, [], [f383, f1763])).
fof(f383, plain, ~ (op1(e10, e11) = op1(e12, e11)), inference(cnf_transformation, [], [f7])).
fof(f14112, plain, (~ spl80_85 | ~ spl80_187 | ~ spl80_286 | ~ spl80_291 | spl80_539), inference(avatar_contradiction_clause, [], [f14111])).
fof(f14111, plain, ($false | (~ spl80_85 | ~ spl80_187 | ~ spl80_286 | ~ spl80_291 | spl80_539)), inference(subsumption_resolution, [], [f14110, f2107])).
fof(f2107, plain, ((e21 = op2(e22, e23)) | ~ spl80_187), inference(avatar_component_clause, [], [f2105])).
fof(f2105, plain, (spl80_187 <=> (e21 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl80_187])])).
fof(f14110, plain, (~ (e21 = op2(e22, e23)) | (~ spl80_85 | ~ spl80_286 | ~ spl80_291 | spl80_539)), inference(forward_demodulation, [], [f14109, f759])).
fof(f759, plain, (e21 = h18(e14)), inference(cnf_transformation, [], [f33])).
fof(f14109, plain, (~ (op2(e22, e23) = h18(e14)) | (~ spl80_85 | ~ spl80_286 | ~ spl80_291 | spl80_539)), inference(forward_demodulation, [], [f14108, f2829])).
fof(f14108, plain, (~ (h18(e14) = op2(h8(e13), e23)) | (~ spl80_85 | ~ spl80_286 | spl80_539)), inference(forward_demodulation, [], [f14107, f2799])).
fof(f2799, plain, ((e23 = h8(e11)) | ~ spl80_286), inference(avatar_component_clause, [], [f2798])).
fof(f2798, plain, (spl80_286 <=> (e23 = h8(e11))), introduced(avatar_definition, [new_symbols(naming, [spl80_286])])).
fof(f14107, plain, (~ (h18(e14) = op2(h8(e13), h8(e11))) | (~ spl80_85 | spl80_539)), inference(forward_demodulation, [], [f4569, f1628])).
fof(f1628, plain, ((e14 = op1(e11, e13)) | ~ spl80_85), inference(avatar_component_clause, [], [f1626])).
fof(f1626, plain, (spl80_85 <=> (e14 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl80_85])])).
fof(f4569, plain, (~ (op2(h8(e13), h8(e11)) = h18(op1(e11, e13))) | spl80_539), inference(avatar_component_clause, [], [f4567])).
fof(f4567, plain, (spl80_539 <=> (op2(h8(e13), h8(e11)) = h18(op1(e11, e13)))), introduced(avatar_definition, [new_symbols(naming, [spl80_539])])).
fof(f14018, plain, (~ spl80_73 | ~ spl80_155 | ~ spl80_296 | spl80_537), inference(avatar_contradiction_clause, [], [f14017])).
fof(f14017, plain, ($false | (~ spl80_73 | ~ spl80_155 | ~ spl80_296 | spl80_537)), inference(subsumption_resolution, [], [f14016, f1972])).
fof(f1972, plain, ((e24 = op2(e24, e20)) | ~ spl80_155), inference(avatar_component_clause, [], [f1970])).
fof(f1970, plain, (spl80_155 <=> (e24 = op2(e24, e20))), introduced(avatar_definition, [new_symbols(naming, [spl80_155])])).
fof(f14016, plain, (~ (e24 = op2(e24, e20)) | (~ spl80_73 | ~ spl80_296 | spl80_537)), inference(forward_demodulation, [], [f14015, f758])).
fof(f758, plain, (e24 = h18(e12)), inference(cnf_transformation, [], [f33])).
fof(f14015, plain, (~ (op2(e24, e20) = h18(e12)) | (~ spl80_73 | ~ spl80_296 | spl80_537)), inference(forward_demodulation, [], [f14014, f1578])).
fof(f1578, plain, ((e12 = op1(e12, e10)) | ~ spl80_73), inference(avatar_component_clause, [], [f1576])).
fof(f1576, plain, (spl80_73 <=> (e12 = op1(e12, e10))), introduced(avatar_definition, [new_symbols(naming, [spl80_73])])).
fof(f14014, plain, (~ (op2(e24, e20) = h18(op1(e12, e10))) | (~ spl80_296 | spl80_537)), inference(forward_demodulation, [], [f4561, f2858])).
fof(f4561, plain, (~ (h18(op1(e12, e10)) = op2(e24, h18(e10))) | spl80_537), inference(avatar_component_clause, [], [f4559])).
fof(f4559, plain, (spl80_537 <=> (h18(op1(e12, e10)) = op2(e24, h18(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl80_537])])).
fof(f13926, plain, (~ spl80_70 | ~ spl80_142 | ~ spl80_291 | spl80_536), inference(avatar_contradiction_clause, [], [f13925])).
fof(f13925, plain, ($false | (~ spl80_70 | ~ spl80_142 | ~ spl80_291 | spl80_536)), inference(subsumption_resolution, [], [f13924, f1918])).
fof(f1918, plain, ((e21 = op2(e24, e22)) | ~ spl80_142), inference(avatar_component_clause, [], [f1916])).
fof(f1916, plain, (spl80_142 <=> (e21 = op2(e24, e22))), introduced(avatar_definition, [new_symbols(naming, [spl80_142])])).
fof(f13924, plain, (~ (e21 = op2(e24, e22)) | (~ spl80_70 | ~ spl80_291 | spl80_536)), inference(forward_demodulation, [], [f13923, f759])).
fof(f13923, plain, (~ (op2(e24, e22) = h18(e14)) | (~ spl80_70 | ~ spl80_291 | spl80_536)), inference(forward_demodulation, [], [f13922, f1565])).
fof(f1565, plain, ((e14 = op1(e12, e11)) | ~ spl80_70), inference(avatar_component_clause, [], [f1563])).
fof(f1563, plain, (spl80_70 <=> (e14 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl80_70])])).
fof(f13922, plain, (~ (op2(e24, e22) = h18(op1(e12, e11))) | (~ spl80_291 | spl80_536)), inference(forward_demodulation, [], [f4557, f2829])).
fof(f4557, plain, (~ (h18(op1(e12, e11)) = op2(e24, h8(e13))) | spl80_536), inference(avatar_component_clause, [], [f4555])).
fof(f4555, plain, (spl80_536 <=> (h18(op1(e12, e11)) = op2(e24, h8(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl80_536])])).
fof(f13832, plain, (~ spl80_117 | ~ spl80_243 | ~ spl80_291 | ~ spl80_296 | spl80_546), inference(avatar_contradiction_clause, [], [f13831])).
fof(f13831, plain, ($false | (~ spl80_117 | ~ spl80_243 | ~ spl80_291 | ~ spl80_296 | spl80_546)), inference(subsumption_resolution, [], [f13830, f2342])).
fof(f2342, plain, ((e22 = op2(e20, e22)) | ~ spl80_243), inference(avatar_component_clause, [], [f2340])).
fof(f2340, plain, (spl80_243 <=> (e22 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl80_243])])).
fof(f13830, plain, (~ (e22 = op2(e20, e22)) | (~ spl80_117 | ~ spl80_291 | ~ spl80_296 | spl80_546)), inference(forward_demodulation, [], [f13829, f11544])).
fof(f13829, plain, (~ (op2(e20, e22) = h18(e11)) | (~ spl80_117 | ~ spl80_291 | ~ spl80_296 | spl80_546)), inference(forward_demodulation, [], [f13828, f1763])).
fof(f13828, plain, (~ (op2(e20, e22) = h18(op1(e10, e11))) | (~ spl80_291 | ~ spl80_296 | spl80_546)), inference(forward_demodulation, [], [f13827, f2858])).
fof(f13827, plain, (~ (h18(op1(e10, e11)) = op2(h18(e10), e22)) | (~ spl80_291 | spl80_546)), inference(forward_demodulation, [], [f4597, f2829])).
fof(f4597, plain, (~ (h18(op1(e10, e11)) = op2(h18(e10), h8(e13))) | spl80_546), inference(avatar_component_clause, [], [f4595])).
fof(f4595, plain, (spl80_546 <=> (h18(op1(e10, e11)) = op2(h18(e10), h8(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl80_546])])).
fof(f13748, plain, (~ spl80_61 | ~ spl80_131 | ~ spl80_296 | spl80_535), inference(avatar_contradiction_clause, [], [f13747])).
fof(f13747, plain, ($false | (~ spl80_61 | ~ spl80_131 | ~ spl80_296 | spl80_535)), inference(subsumption_resolution, [], [f13746, f1872])).
fof(f1872, plain, ((e20 = op2(e24, e24)) | ~ spl80_131), inference(avatar_component_clause, [], [f1870])).
fof(f1870, plain, (spl80_131 <=> (e20 = op2(e24, e24))), introduced(avatar_definition, [new_symbols(naming, [spl80_131])])).
fof(f13746, plain, (~ (e20 = op2(e24, e24)) | (~ spl80_61 | ~ spl80_296 | spl80_535)), inference(forward_demodulation, [], [f13745, f2858])).
fof(f13745, plain, (~ (op2(e24, e24) = h18(e10)) | (~ spl80_61 | spl80_535)), inference(forward_demodulation, [], [f4553, f1528])).
fof(f1528, plain, ((e10 = op1(e12, e12)) | ~ spl80_61), inference(avatar_component_clause, [], [f1526])).
fof(f1526, plain, (spl80_61 <=> (e10 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl80_61])])).
fof(f4553, plain, (~ (op2(e24, e24) = h18(op1(e12, e12))) | spl80_535), inference(avatar_component_clause, [], [f4551])).
fof(f4551, plain, (spl80_535 <=> (op2(e24, e24) = h18(op1(e12, e12)))), introduced(avatar_definition, [new_symbols(naming, [spl80_535])])).
fof(f13659, plain, (~ spl80_57 | ~ spl80_138 | ~ spl80_286 | ~ spl80_291 | spl80_534), inference(avatar_contradiction_clause, [], [f13658])).
fof(f13658, plain, ($false | (~ spl80_57 | ~ spl80_138 | ~ spl80_286 | ~ spl80_291 | spl80_534)), inference(subsumption_resolution, [], [f13657, f1901])).
fof(f1901, plain, ((e22 = op2(e24, e23)) | ~ spl80_138), inference(avatar_component_clause, [], [f1899])).
fof(f1899, plain, (spl80_138 <=> (e22 = op2(e24, e23))), introduced(avatar_definition, [new_symbols(naming, [spl80_138])])).
fof(f13657, plain, (~ (e22 = op2(e24, e23)) | (~ spl80_57 | ~ spl80_286 | ~ spl80_291 | spl80_534)), inference(forward_demodulation, [], [f13656, f11544])).
fof(f13656, plain, (~ (op2(e24, e23) = h18(e11)) | (~ spl80_57 | ~ spl80_286 | spl80_534)), inference(forward_demodulation, [], [f13655, f1511])).
fof(f1511, plain, ((e11 = op1(e12, e13)) | ~ spl80_57), inference(avatar_component_clause, [], [f1509])).
fof(f1509, plain, (spl80_57 <=> (e11 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl80_57])])).
fof(f13655, plain, (~ (op2(e24, e23) = h18(op1(e12, e13))) | (~ spl80_286 | spl80_534)), inference(forward_demodulation, [], [f4549, f2799])).
fof(f4549, plain, (~ (h18(op1(e12, e13)) = op2(e24, h8(e11))) | spl80_534), inference(avatar_component_clause, [], [f4547])).
fof(f4547, plain, (spl80_534 <=> (h18(op1(e12, e13)) = op2(e24, h8(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl80_534])])).
fof(f13568, plain, (~ spl80_43 | ~ spl80_170 | ~ spl80_286 | ~ spl80_291 | spl80_531), inference(avatar_contradiction_clause, [], [f13567])).
fof(f13567, plain, ($false | (~ spl80_43 | ~ spl80_170 | ~ spl80_286 | ~ spl80_291 | spl80_531)), inference(subsumption_resolution, [], [f13566, f2035])).
fof(f2035, plain, ((e24 = op2(e23, e22)) | ~ spl80_170), inference(avatar_component_clause, [], [f2033])).
fof(f2033, plain, (spl80_170 <=> (e24 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl80_170])])).
fof(f13566, plain, (~ (e24 = op2(e23, e22)) | (~ spl80_43 | ~ spl80_286 | ~ spl80_291 | spl80_531)), inference(forward_demodulation, [], [f13565, f758])).
fof(f13565, plain, (~ (op2(e23, e22) = h18(e12)) | (~ spl80_43 | ~ spl80_286 | ~ spl80_291 | spl80_531)), inference(forward_demodulation, [], [f13564, f2799])).
fof(f13564, plain, (~ (h18(e12) = op2(h8(e11), e22)) | (~ spl80_43 | ~ spl80_291 | spl80_531)), inference(forward_demodulation, [], [f13563, f2829])).
fof(f13563, plain, (~ (h18(e12) = op2(h8(e11), h8(e13))) | (~ spl80_43 | spl80_531)), inference(forward_demodulation, [], [f4537, f1452])).
fof(f1452, plain, ((e12 = op1(e13, e11)) | ~ spl80_43), inference(avatar_component_clause, [], [f1450])).
fof(f1450, plain, (spl80_43 <=> (e12 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl80_43])])).
fof(f4537, plain, (~ (op2(h8(e11), h8(e13)) = h18(op1(e13, e11))) | spl80_531), inference(avatar_component_clause, [], [f4535])).
fof(f4535, plain, (spl80_531 <=> (op2(h8(e11), h8(e13)) = h18(op1(e13, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl80_531])])).
fof(f13486, plain, (~ spl80_40 | ~ spl80_157 | ~ spl80_286 | spl80_530), inference(avatar_contradiction_clause, [], [f13485])).
fof(f13485, plain, ($false | (~ spl80_40 | ~ spl80_157 | ~ spl80_286 | spl80_530)), inference(subsumption_resolution, [], [f13484, f1981])).
fof(f1981, plain, ((e21 = op2(e23, e24)) | ~ spl80_157), inference(avatar_component_clause, [], [f1979])).
fof(f1979, plain, (spl80_157 <=> (e21 = op2(e23, e24))), introduced(avatar_definition, [new_symbols(naming, [spl80_157])])).
fof(f13484, plain, (~ (e21 = op2(e23, e24)) | (~ spl80_40 | ~ spl80_286 | spl80_530)), inference(forward_demodulation, [], [f13483, f759])).
fof(f13483, plain, (~ (op2(e23, e24) = h18(e14)) | (~ spl80_40 | ~ spl80_286 | spl80_530)), inference(forward_demodulation, [], [f13482, f1439])).
fof(f1439, plain, ((e14 = op1(e13, e12)) | ~ spl80_40), inference(avatar_component_clause, [], [f1437])).
fof(f1437, plain, (spl80_40 <=> (e14 = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl80_40])])).
fof(f13482, plain, (~ (op2(e23, e24) = h18(op1(e13, e12))) | (~ spl80_286 | spl80_530)), inference(forward_demodulation, [], [f4533, f2799])).
fof(f4533, plain, (~ (h18(op1(e13, e12)) = op2(h8(e11), e24)) | spl80_530), inference(avatar_component_clause, [], [f4531])).
fof(f4531, plain, (spl80_530 <=> (h18(op1(e13, e12)) = op2(h8(e11), e24))), introduced(avatar_definition, [new_symbols(naming, [spl80_530])])).
fof(f13397, plain, (~ spl80_113 | ~ spl80_235 | ~ spl80_296 | spl80_545), inference(avatar_contradiction_clause, [], [f13396])).
fof(f13396, plain, ($false | (~ spl80_113 | ~ spl80_235 | ~ spl80_296 | spl80_545)), inference(subsumption_resolution, [], [f13395, f2308])).
fof(f2308, plain, ((e24 = op2(e20, e24)) | ~ spl80_235), inference(avatar_component_clause, [], [f2306])).
fof(f2306, plain, (spl80_235 <=> (e24 = op2(e20, e24))), introduced(avatar_definition, [new_symbols(naming, [spl80_235])])).
fof(f13395, plain, (~ (e24 = op2(e20, e24)) | (~ spl80_113 | ~ spl80_296 | spl80_545)), inference(forward_demodulation, [], [f13394, f758])).
fof(f13394, plain, (~ (op2(e20, e24) = h18(e12)) | (~ spl80_113 | ~ spl80_296 | spl80_545)), inference(forward_demodulation, [], [f13393, f1746])).
fof(f13393, plain, (~ (op2(e20, e24) = h18(op1(e10, e12))) | (~ spl80_296 | spl80_545)), inference(forward_demodulation, [], [f4593, f2858])).
fof(f4593, plain, (~ (h18(op1(e10, e12)) = op2(h18(e10), e24)) | spl80_545), inference(avatar_component_clause, [], [f4591])).
fof(f4591, plain, (spl80_545 <=> (h18(op1(e10, e12)) = op2(h18(e10), e24))), introduced(avatar_definition, [new_symbols(naming, [spl80_545])])).
fof(f13306, plain, (~ spl80_105 | ~ spl80_247 | ~ spl80_296 | spl80_543), inference(avatar_contradiction_clause, [], [f13305])).
fof(f13305, plain, ($false | (~ spl80_105 | ~ spl80_247 | ~ spl80_296 | spl80_543)), inference(subsumption_resolution, [], [f13304, f2359])).
fof(f2359, plain, ((e21 = op2(e20, e21)) | ~ spl80_247), inference(avatar_component_clause, [], [f2357])).
fof(f2357, plain, (spl80_247 <=> (e21 = op2(e20, e21))), introduced(avatar_definition, [new_symbols(naming, [spl80_247])])).
fof(f13304, plain, (~ (e21 = op2(e20, e21)) | (~ spl80_105 | ~ spl80_296 | spl80_543)), inference(forward_demodulation, [], [f13303, f759])).
fof(f13303, plain, (~ (op2(e20, e21) = h18(e14)) | (~ spl80_105 | ~ spl80_296 | spl80_543)), inference(forward_demodulation, [], [f13302, f1712])).
fof(f1712, plain, ((e14 = op1(e10, e14)) | ~ spl80_105), inference(avatar_component_clause, [], [f1710])).
fof(f1710, plain, (spl80_105 <=> (e14 = op1(e10, e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_105])])).
fof(f13302, plain, (~ (op2(e20, e21) = h18(op1(e10, e14))) | (~ spl80_296 | spl80_543)), inference(forward_demodulation, [], [f4585, f2858])).
fof(f4585, plain, (~ (h18(op1(e10, e14)) = op2(h18(e10), e21)) | spl80_543), inference(avatar_component_clause, [], [f4583])).
fof(f4583, plain, (spl80_543 <=> (h18(op1(e10, e14)) = op2(h18(e10), e21))), introduced(avatar_definition, [new_symbols(naming, [spl80_543])])).
fof(f13214, plain, (~ spl80_78 | ~ spl80_200 | ~ spl80_291 | spl80_538), inference(avatar_contradiction_clause, [], [f13213])).
fof(f13213, plain, ($false | (~ spl80_78 | ~ spl80_200 | ~ spl80_291 | spl80_538)), inference(subsumption_resolution, [], [f13212, f2161])).
fof(f2161, plain, ((e24 = op2(e22, e21)) | ~ spl80_200), inference(avatar_component_clause, [], [f2159])).
fof(f2159, plain, (spl80_200 <=> (e24 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl80_200])])).
fof(f13212, plain, (~ (e24 = op2(e22, e21)) | (~ spl80_78 | ~ spl80_291 | spl80_538)), inference(forward_demodulation, [], [f13211, f758])).
fof(f13211, plain, (~ (op2(e22, e21) = h18(e12)) | (~ spl80_78 | ~ spl80_291 | spl80_538)), inference(forward_demodulation, [], [f13210, f1599])).
fof(f1599, plain, ((e12 = op1(e11, e14)) | ~ spl80_78), inference(avatar_component_clause, [], [f1597])).
fof(f1597, plain, (spl80_78 <=> (e12 = op1(e11, e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_78])])).
fof(f13210, plain, (~ (op2(e22, e21) = h18(op1(e11, e14))) | (~ spl80_291 | spl80_538)), inference(forward_demodulation, [], [f4565, f2829])).
fof(f4565, plain, (~ (h18(op1(e11, e14)) = op2(h8(e13), e21)) | spl80_538), inference(avatar_component_clause, [], [f4563])).
fof(f4563, plain, (spl80_538 <=> (h18(op1(e11, e14)) = op2(h8(e13), e21))), introduced(avatar_definition, [new_symbols(naming, [spl80_538])])).
fof(f13119, plain, (~ spl80_54 | ~ spl80_286 | spl80_533), inference(avatar_contradiction_clause, [], [f13118])).
fof(f13118, plain, ($false | (~ spl80_54 | ~ spl80_286 | spl80_533)), inference(subsumption_resolution, [], [f13117, f2799])).
fof(f13117, plain, (~ (e23 = h8(e11)) | (~ spl80_54 | ~ spl80_286 | spl80_533)), inference(forward_demodulation, [], [f13116, f12007])).
fof(f12007, plain, ((e23 = h18(e13)) | ~ spl80_286), inference(backward_demodulation, [], [f2642, f2799])).
fof(f2642, plain, (h8(e11) = h18(e13)), inference(forward_demodulation, [], [f762, f711])).
fof(f711, plain, (op2(e24, e21) = h8(e11)), inference(cnf_transformation, [], [f23])).
fof(f762, plain, (op2(e24, e21) = h18(e13)), inference(cnf_transformation, [], [f33])).
fof(f13116, plain, (~ (h8(e11) = h18(e13)) | (~ spl80_54 | spl80_533)), inference(forward_demodulation, [], [f4545, f1498])).
fof(f1498, plain, ((e13 = op1(e12, e14)) | ~ spl80_54), inference(avatar_component_clause, [], [f1496])).
fof(f1496, plain, (spl80_54 <=> (e13 = op1(e12, e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_54])])).
fof(f4545, plain, (~ (h8(e11) = h18(op1(e12, e14))) | spl80_533), inference(avatar_component_clause, [], [f4543])).
fof(f4543, plain, (spl80_533 <=> (h8(e11) = h18(op1(e12, e14)))), introduced(avatar_definition, [new_symbols(naming, [spl80_533])])).
fof(f13040, plain, (~ spl80_49 | ~ spl80_179 | ~ spl80_286 | ~ spl80_296 | spl80_532), inference(avatar_contradiction_clause, [], [f13039])).
fof(f13039, plain, ($false | (~ spl80_49 | ~ spl80_179 | ~ spl80_286 | ~ spl80_296 | spl80_532)), inference(subsumption_resolution, [], [f13038, f2073])).
fof(f2073, plain, ((e23 = op2(e23, e20)) | ~ spl80_179), inference(avatar_component_clause, [], [f2071])).
fof(f2071, plain, (spl80_179 <=> (e23 = op2(e23, e20))), introduced(avatar_definition, [new_symbols(naming, [spl80_179])])).
fof(f13038, plain, (~ (e23 = op2(e23, e20)) | (~ spl80_49 | ~ spl80_286 | ~ spl80_296 | spl80_532)), inference(forward_demodulation, [], [f13037, f12007])).
fof(f13037, plain, (~ (op2(e23, e20) = h18(e13)) | (~ spl80_49 | ~ spl80_286 | ~ spl80_296 | spl80_532)), inference(forward_demodulation, [], [f13036, f1477])).
fof(f13036, plain, (~ (op2(e23, e20) = h18(op1(e13, e10))) | (~ spl80_286 | ~ spl80_296 | spl80_532)), inference(forward_demodulation, [], [f13035, f2799])).
fof(f13035, plain, (~ (h18(op1(e13, e10)) = op2(h8(e11), e20)) | (~ spl80_296 | spl80_532)), inference(forward_demodulation, [], [f4541, f2858])).
fof(f4541, plain, (~ (h18(op1(e13, e10)) = op2(h8(e11), h18(e10))) | spl80_532), inference(avatar_component_clause, [], [f4539])).
fof(f4539, plain, (spl80_532 <=> (h18(op1(e13, e10)) = op2(h8(e11), h18(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl80_532])])).
fof(f12956, plain, (~ spl80_31 | ~ spl80_296 | ~ spl80_402 | spl80_529), inference(avatar_contradiction_clause, [], [f12955])).
fof(f12955, plain, ($false | (~ spl80_31 | ~ spl80_296 | ~ spl80_402 | spl80_529)), inference(subsumption_resolution, [], [f12954, f3500])).
fof(f3500, plain, ((e20 = h8(e10)) | ~ spl80_402), inference(avatar_component_clause, [], [f3499])).
fof(f3499, plain, (spl80_402 <=> (e20 = h8(e10))), introduced(avatar_definition, [new_symbols(naming, [spl80_402])])).
fof(f12954, plain, (~ (e20 = h8(e10)) | (~ spl80_31 | ~ spl80_296 | spl80_529)), inference(forward_demodulation, [], [f12953, f2858])).
fof(f12953, plain, (~ (h8(e10) = h18(e10)) | (~ spl80_31 | spl80_529)), inference(forward_demodulation, [], [f4529, f1402])).
fof(f1402, plain, ((e10 = op1(e13, e13)) | ~ spl80_31), inference(avatar_component_clause, [], [f1400])).
fof(f1400, plain, (spl80_31 <=> (e10 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl80_31])])).
fof(f4529, plain, (~ (h8(e10) = h18(op1(e13, e13))) | spl80_529), inference(avatar_component_clause, [], [f4527])).
fof(f4527, plain, (spl80_529 <=> (h8(e10) = h18(op1(e13, e13)))), introduced(avatar_definition, [new_symbols(naming, [spl80_529])])).
fof(f12867, plain, (~ spl80_27 | ~ spl80_173 | ~ spl80_286 | ~ spl80_291 | spl80_528), inference(avatar_contradiction_clause, [], [f12866])).
fof(f12866, plain, ($false | (~ spl80_27 | ~ spl80_173 | ~ spl80_286 | ~ spl80_291 | spl80_528)), inference(subsumption_resolution, [], [f12865, f2048])).
fof(f2048, plain, ((e22 = op2(e23, e21)) | ~ spl80_173), inference(avatar_component_clause, [], [f2046])).
fof(f2046, plain, (spl80_173 <=> (e22 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl80_173])])).
fof(f12865, plain, (~ (e22 = op2(e23, e21)) | (~ spl80_27 | ~ spl80_286 | ~ spl80_291 | spl80_528)), inference(forward_demodulation, [], [f12864, f11544])).
fof(f12864, plain, (~ (op2(e23, e21) = h18(e11)) | (~ spl80_27 | ~ spl80_286 | spl80_528)), inference(forward_demodulation, [], [f12863, f1385])).
fof(f12863, plain, (~ (op2(e23, e21) = h18(op1(e13, e14))) | (~ spl80_286 | spl80_528)), inference(forward_demodulation, [], [f4525, f2799])).
fof(f4525, plain, (~ (h18(op1(e13, e14)) = op2(h8(e11), e21)) | spl80_528), inference(avatar_component_clause, [], [f4523])).
fof(f4523, plain, (spl80_528 <=> (h18(op1(e13, e14)) = op2(h8(e11), e21))), introduced(avatar_definition, [new_symbols(naming, [spl80_528])])).
fof(f12769, plain, (~ spl80_25 | ~ spl80_227 | ~ spl80_296 | spl80_527), inference(avatar_contradiction_clause, [], [f12768])).
fof(f12768, plain, ($false | (~ spl80_25 | ~ spl80_227 | ~ spl80_296 | spl80_527)), inference(subsumption_resolution, [], [f12767, f2275])).
fof(f2275, plain, ((e21 = op2(e21, e20)) | ~ spl80_227), inference(avatar_component_clause, [], [f2273])).
fof(f2273, plain, (spl80_227 <=> (e21 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl80_227])])).
fof(f12767, plain, (~ (e21 = op2(e21, e20)) | (~ spl80_25 | ~ spl80_296 | spl80_527)), inference(forward_demodulation, [], [f12766, f759])).
fof(f12766, plain, (~ (op2(e21, e20) = h18(e14)) | (~ spl80_25 | ~ spl80_296 | spl80_527)), inference(forward_demodulation, [], [f12765, f1376])).
fof(f1376, plain, ((e14 = op1(e14, e10)) | ~ spl80_25), inference(avatar_component_clause, [], [f1374])).
fof(f1374, plain, (spl80_25 <=> (e14 = op1(e14, e10))), introduced(avatar_definition, [new_symbols(naming, [spl80_25])])).
fof(f12765, plain, (~ (op2(e21, e20) = h18(op1(e14, e10))) | (~ spl80_296 | spl80_527)), inference(forward_demodulation, [], [f4521, f2858])).
fof(f4521, plain, (~ (h18(op1(e14, e10)) = op2(e21, h18(e10))) | spl80_527), inference(avatar_component_clause, [], [f4519])).
fof(f4519, plain, (spl80_527 <=> (h18(op1(e14, e10)) = op2(e21, h18(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl80_527])])).
fof(f12660, plain, (~ spl80_19 | ~ spl80_219 | ~ spl80_286 | ~ spl80_291 | spl80_526), inference(avatar_contradiction_clause, [], [f12659])).
fof(f12659, plain, ($false | (~ spl80_19 | ~ spl80_219 | ~ spl80_286 | ~ spl80_291 | spl80_526)), inference(subsumption_resolution, [], [f12658, f2241])).
fof(f2241, plain, ((e23 = op2(e21, e22)) | ~ spl80_219), inference(avatar_component_clause, [], [f2239])).
fof(f2239, plain, (spl80_219 <=> (e23 = op2(e21, e22))), introduced(avatar_definition, [new_symbols(naming, [spl80_219])])).
fof(f12658, plain, (~ (e23 = op2(e21, e22)) | (~ spl80_19 | ~ spl80_286 | ~ spl80_291 | spl80_526)), inference(forward_demodulation, [], [f12657, f12007])).
fof(f12657, plain, (~ (op2(e21, e22) = h18(e13)) | (~ spl80_19 | ~ spl80_291 | spl80_526)), inference(forward_demodulation, [], [f12656, f1351])).
fof(f1351, plain, ((e13 = op1(e14, e11)) | ~ spl80_19), inference(avatar_component_clause, [], [f1349])).
fof(f1349, plain, (spl80_19 <=> (e13 = op1(e14, e11))), introduced(avatar_definition, [new_symbols(naming, [spl80_19])])).
fof(f12656, plain, (~ (op2(e21, e22) = h18(op1(e14, e11))) | (~ spl80_291 | spl80_526)), inference(forward_demodulation, [], [f4517, f2829])).
fof(f4517, plain, (~ (h18(op1(e14, e11)) = op2(e21, h8(e13))) | spl80_526), inference(avatar_component_clause, [], [f4515])).
fof(f4515, plain, (spl80_526 <=> (h18(op1(e14, e11)) = op2(e21, h8(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl80_526])])).
fof(f12532, plain, (~ spl80_12 | ~ spl80_291 | spl80_525), inference(avatar_contradiction_clause, [], [f12531])).
fof(f12531, plain, ($false | (~ spl80_12 | ~ spl80_291 | spl80_525)), inference(subsumption_resolution, [], [f12530, f2829])).
fof(f12530, plain, (~ (e22 = h8(e13)) | (~ spl80_12 | ~ spl80_291 | spl80_525)), inference(forward_demodulation, [], [f12529, f11544])).
fof(f12529, plain, (~ (h8(e13) = h18(e11)) | (~ spl80_12 | spl80_525)), inference(forward_demodulation, [], [f4513, f1322])).
fof(f1322, plain, ((e11 = op1(e14, e12)) | ~ spl80_12), inference(avatar_component_clause, [], [f1320])).
fof(f1320, plain, (spl80_12 <=> (e11 = op1(e14, e12))), introduced(avatar_definition, [new_symbols(naming, [spl80_12])])).
fof(f4513, plain, (~ (h8(e13) = h18(op1(e14, e12))) | spl80_525), inference(avatar_component_clause, [], [f4511])).
fof(f4511, plain, (spl80_525 <=> (h8(e13) = h18(op1(e14, e12)))), introduced(avatar_definition, [new_symbols(naming, [spl80_525])])).
fof(f12403, plain, (~ spl80_8 | ~ spl80_215 | ~ spl80_286 | spl80_524), inference(avatar_contradiction_clause, [], [f12402])).
fof(f12402, plain, ($false | (~ spl80_8 | ~ spl80_215 | ~ spl80_286 | spl80_524)), inference(subsumption_resolution, [], [f12401, f2224])).
fof(f2224, plain, ((e24 = op2(e21, e23)) | ~ spl80_215), inference(avatar_component_clause, [], [f2222])).
fof(f2222, plain, (spl80_215 <=> (e24 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl80_215])])).
fof(f12401, plain, (~ (e24 = op2(e21, e23)) | (~ spl80_8 | ~ spl80_286 | spl80_524)), inference(forward_demodulation, [], [f12400, f758])).
fof(f12400, plain, (~ (op2(e21, e23) = h18(e12)) | (~ spl80_8 | ~ spl80_286 | spl80_524)), inference(forward_demodulation, [], [f12399, f1305])).
fof(f1305, plain, ((e12 = op1(e14, e13)) | ~ spl80_8), inference(avatar_component_clause, [], [f1303])).
fof(f1303, plain, (spl80_8 <=> (e12 = op1(e14, e13))), introduced(avatar_definition, [new_symbols(naming, [spl80_8])])).
fof(f12399, plain, (~ (op2(e21, e23) = h18(op1(e14, e13))) | (~ spl80_286 | spl80_524)), inference(forward_demodulation, [], [f4509, f2799])).
fof(f4509, plain, (~ (h18(op1(e14, e13)) = op2(e21, h8(e11))) | spl80_524), inference(avatar_component_clause, [], [f4507])).
fof(f4507, plain, (spl80_524 <=> (h18(op1(e14, e13)) = op2(e21, h8(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl80_524])])).
fof(f12272, plain, (~ spl80_109 | ~ spl80_239 | ~ spl80_286 | ~ spl80_296 | spl80_544), inference(avatar_contradiction_clause, [], [f12271])).
fof(f12271, plain, ($false | (~ spl80_109 | ~ spl80_239 | ~ spl80_286 | ~ spl80_296 | spl80_544)), inference(subsumption_resolution, [], [f12270, f2325])).
fof(f2325, plain, ((e23 = op2(e20, e23)) | ~ spl80_239), inference(avatar_component_clause, [], [f2323])).
fof(f2323, plain, (spl80_239 <=> (e23 = op2(e20, e23))), introduced(avatar_definition, [new_symbols(naming, [spl80_239])])).
fof(f12270, plain, (~ (e23 = op2(e20, e23)) | (~ spl80_109 | ~ spl80_286 | ~ spl80_296 | spl80_544)), inference(forward_demodulation, [], [f12269, f12007])).
fof(f12269, plain, (~ (op2(e20, e23) = h18(e13)) | (~ spl80_109 | ~ spl80_286 | ~ spl80_296 | spl80_544)), inference(forward_demodulation, [], [f12268, f1729])).
fof(f1729, plain, ((e13 = op1(e10, e13)) | ~ spl80_109), inference(avatar_component_clause, [], [f1727])).
fof(f1727, plain, (spl80_109 <=> (e13 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl80_109])])).
fof(f12268, plain, (~ (op2(e20, e23) = h18(op1(e10, e13))) | (~ spl80_286 | ~ spl80_296 | spl80_544)), inference(forward_demodulation, [], [f12267, f2858])).
fof(f12267, plain, (~ (h18(op1(e10, e13)) = op2(h18(e10), e23)) | (~ spl80_286 | spl80_544)), inference(forward_demodulation, [], [f4589, f2799])).
fof(f4589, plain, (~ (h18(op1(e10, e13)) = op2(h18(e10), h8(e11))) | spl80_544), inference(avatar_component_clause, [], [f4587])).
fof(f4587, plain, (spl80_544 <=> (h18(op1(e10, e13)) = op2(h18(e10), h8(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl80_544])])).
fof(f12142, plain, (~ spl80_91 | spl80_541), inference(avatar_contradiction_clause, [], [f12141])).
fof(f12141, plain, ($false | (~ spl80_91 | spl80_541)), inference(trivial_inequality_removal, [], [f12140])).
fof(f12140, plain, (~ (h18(e10) = h18(e10)) | (~ spl80_91 | spl80_541)), inference(forward_demodulation, [], [f4577, f1654])).
fof(f1654, plain, ((e10 = op1(e11, e11)) | ~ spl80_91), inference(avatar_component_clause, [], [f1652])).
fof(f1652, plain, (spl80_91 <=> (e10 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl80_91])])).
fof(f4577, plain, (~ (h18(e10) = h18(op1(e11, e11))) | spl80_541), inference(avatar_component_clause, [], [f4575])).
fof(f4575, plain, (spl80_541 <=> (h18(e10) = h18(op1(e11, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl80_541])])).
fof(f11974, plain, (~ spl80_286 | ~ spl80_89 | ~ spl80_184 | ~ spl80_291 | spl80_540), inference(avatar_split_clause, [], [f11973, f4571, f2828, f2092, f1643, f2798])).
fof(f1643, plain, (spl80_89 <=> (e13 = op1(e11, e12))), introduced(avatar_definition, [new_symbols(naming, [spl80_89])])).
fof(f2092, plain, (spl80_184 <=> (e23 = op2(e22, e24))), introduced(avatar_definition, [new_symbols(naming, [spl80_184])])).
fof(f4571, plain, (spl80_540 <=> (h18(op1(e11, e12)) = op2(h8(e13), e24))), introduced(avatar_definition, [new_symbols(naming, [spl80_540])])).
fof(f11973, plain, (~ (e23 = h8(e11)) | (~ spl80_89 | ~ spl80_184 | ~ spl80_291 | spl80_540)), inference(forward_demodulation, [], [f11972, f2094])).
fof(f2094, plain, ((e23 = op2(e22, e24)) | ~ spl80_184), inference(avatar_component_clause, [], [f2092])).
fof(f11972, plain, (~ (op2(e22, e24) = h8(e11)) | (~ spl80_89 | ~ spl80_291 | spl80_540)), inference(forward_demodulation, [], [f11971, f2642])).
fof(f11971, plain, (~ (op2(e22, e24) = h18(e13)) | (~ spl80_89 | ~ spl80_291 | spl80_540)), inference(forward_demodulation, [], [f11970, f1645])).
fof(f1645, plain, ((e13 = op1(e11, e12)) | ~ spl80_89), inference(avatar_component_clause, [], [f1643])).
fof(f11970, plain, (~ (op2(e22, e24) = h18(op1(e11, e12))) | (~ spl80_291 | spl80_540)), inference(forward_demodulation, [], [f4573, f2829])).
fof(f4573, plain, (~ (h18(op1(e11, e12)) = op2(h8(e13), e24)) | spl80_540), inference(avatar_component_clause, [], [f4571])).
fof(f11831, plain, (~ spl80_187 | ~ spl80_190), inference(avatar_contradiction_clause, [], [f11830])).
fof(f11830, plain, ($false | (~ spl80_187 | ~ spl80_190)), inference(subsumption_resolution, [], [f11829, f588])).
fof(f588, plain, ~ (e21 = e24), inference(cnf_transformation, [], [f10])).
fof(f10, plain, (~ (e23 = e24) & ~ (e22 = e24) & ~ (e22 = e23) & ~ (e21 = e24) & ~ (e21 = e23) & ~ (e21 = e22) & ~ (e20 = e24) & ~ (e20 = e23) & ~ (e20 = e22) & ~ (e20 = e21)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG049+1.p', ax10)).
fof(f11829, plain, ((e21 = e24) | (~ spl80_187 | ~ spl80_190)), inference(forward_demodulation, [], [f2119, f2107])).
fof(f2119, plain, ((e24 = op2(e22, e23)) | ~ spl80_190), inference(avatar_component_clause, [], [f2117])).
fof(f2117, plain, (spl80_190 <=> (e24 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl80_190])])).
fof(f11828, plain, (~ spl80_199 | spl80_370), inference(avatar_contradiction_clause, [], [f11827])).
fof(f11827, plain, ($false | (~ spl80_199 | spl80_370)), inference(subsumption_resolution, [], [f11826, f3309])).
fof(f3309, plain, (~ (e23 = h6(e11)) | spl80_370), inference(avatar_component_clause, [], [f3307])).
fof(f3307, plain, (spl80_370 <=> (e23 = h6(e11))), introduced(avatar_definition, [new_symbols(naming, [spl80_370])])).
fof(f11826, plain, ((e23 = h6(e11)) | ~ spl80_199), inference(backward_demodulation, [], [f701, f2157])).
fof(f2157, plain, ((e23 = op2(e22, e21)) | ~ spl80_199), inference(avatar_component_clause, [], [f2155])).
fof(f2155, plain, (spl80_199 <=> (e23 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl80_199])])).
fof(f701, plain, (op2(e22, e21) = h6(e11)), inference(cnf_transformation, [], [f21])).
fof(f21, plain, ((op2(e21, e22) = h6(e13)) & (op2(e22, e21) = h6(e11)) & (h6(e10) = op2(op2(e22, e21), op2(e22, e21))) & (e22 = h6(e14)) & (e21 = h6(e12))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG049+1.p', ax21)).
fof(f11822, plain, (~ spl80_219 | ~ spl80_220), inference(avatar_contradiction_clause, [], [f11821])).
fof(f11821, plain, ($false | (~ spl80_219 | ~ spl80_220)), inference(subsumption_resolution, [], [f11820, f591])).
fof(f591, plain, ~ (e23 = e24), inference(cnf_transformation, [], [f10])).
fof(f11820, plain, ((e23 = e24) | (~ spl80_219 | ~ spl80_220)), inference(forward_demodulation, [], [f2245, f2241])).
fof(f2245, plain, ((e24 = op2(e21, e22)) | ~ spl80_220), inference(avatar_component_clause, [], [f2243])).
fof(f2243, plain, (spl80_220 <=> (e24 = op2(e21, e22))), introduced(avatar_definition, [new_symbols(naming, [spl80_220])])).
fof(f11780, plain, (~ spl80_1 | ~ spl80_221 | ~ spl80_296 | spl80_523), inference(avatar_contradiction_clause, [], [f11779])).
fof(f11779, plain, ($false | (~ spl80_1 | ~ spl80_221 | ~ spl80_296 | spl80_523)), inference(subsumption_resolution, [], [f11778, f2250])).
fof(f2250, plain, ((e20 = op2(e21, e21)) | ~ spl80_221), inference(avatar_component_clause, [], [f2248])).
fof(f2248, plain, (spl80_221 <=> (e20 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl80_221])])).
fof(f11778, plain, (~ (e20 = op2(e21, e21)) | (~ spl80_1 | ~ spl80_296 | spl80_523)), inference(forward_demodulation, [], [f11777, f2858])).
fof(f11777, plain, (~ (op2(e21, e21) = h18(e10)) | (~ spl80_1 | spl80_523)), inference(forward_demodulation, [], [f4505, f1276])).
fof(f1276, plain, ((e10 = op1(e14, e14)) | ~ spl80_1), inference(avatar_component_clause, [], [f1274])).
fof(f1274, plain, (spl80_1 <=> (e10 = op1(e14, e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_1])])).
fof(f4505, plain, (~ (op2(e21, e21) = h18(op1(e14, e14))) | spl80_523), inference(avatar_component_clause, [], [f4503])).
fof(f4503, plain, (spl80_523 <=> (op2(e21, e21) = h18(op1(e14, e14)))), introduced(avatar_definition, [new_symbols(naming, [spl80_523])])).
fof(f11648, plain, (~ spl80_174 | ~ spl80_344), inference(avatar_split_clause, [], [f11459, f3141, f2050])).
fof(f2050, plain, (spl80_174 <=> (e23 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl80_174])])).
fof(f3141, plain, (spl80_344 <=> (e23 = h3(e11))), introduced(avatar_definition, [new_symbols(naming, [spl80_344])])).
fof(f11459, plain, (~ (e23 = op2(e23, e21)) | ~ spl80_344), inference(forward_demodulation, [], [f2549, f3142])).
fof(f3142, plain, ((e23 = h3(e11)) | ~ spl80_344), inference(avatar_component_clause, [], [f3141])).
fof(f2549, plain, ~ (op2(e23, e21) = h3(e11)), inference(backward_demodulation, [], [f552, f686])).
fof(f686, plain, (op2(e23, e20) = h3(e11)), inference(cnf_transformation, [], [f18])).
fof(f18, plain, ((op2(e20, e23) = h3(e13)) & (op2(e23, e20) = h3(e11)) & (h3(e10) = op2(op2(e23, e20), op2(e23, e20))) & (e23 = h3(e14)) & (e20 = h3(e12))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG049+1.p', ax18)).
fof(f552, plain, ~ (op2(e23, e20) = op2(e23, e21)), inference(cnf_transformation, [], [f8])).
fof(f8, plain, (~ (op2(e24, e23) = op2(e24, e24)) & ~ (op2(e24, e22) = op2(e24, e24)) & ~ (op2(e24, e22) = op2(e24, e23)) & ~ (op2(e24, e21) = op2(e24, e24)) & ~ (op2(e24, e21) = op2(e24, e23)) & ~ (op2(e24, e21) = op2(e24, e22)) & ~ (op2(e24, e20) = op2(e24, e24)) & ~ (op2(e24, e20) = op2(e24, e23)) & ~ (op2(e24, e20) = op2(e24, e22)) & ~ (op2(e24, e20) = op2(e24, e21)) & ~ (op2(e23, e23) = op2(e23, e24)) & ~ (op2(e23, e22) = op2(e23, e24)) & ~ (op2(e23, e22) = op2(e23, e23)) & ~ (op2(e23, e21) = op2(e23, e24)) & ~ (op2(e23, e21) = op2(e23, e23)) & ~ (op2(e23, e21) = op2(e23, e22)) & ~ (op2(e23, e20) = op2(e23, e24)) & ~ (op2(e23, e20) = op2(e23, e23)) & ~ (op2(e23, e20) = op2(e23, e22)) & ~ (op2(e23, e20) = op2(e23, e21)) & ~ (op2(e22, e23) = op2(e22, e24)) & ~ (op2(e22, e22) = op2(e22, e24)) & ~ (op2(e22, e22) = op2(e22, e23)) & ~ (op2(e22, e21) = op2(e22, e24)) & ~ (op2(e22, e21) = op2(e22, e23)) & ~ (op2(e22, e21) = op2(e22, e22)) & ~ (op2(e22, e20) = op2(e22, e24)) & ~ (op2(e22, e20) = op2(e22, e23)) & ~ (op2(e22, e20) = op2(e22, e22)) & ~ (op2(e22, e20) = op2(e22, e21)) & ~ (op2(e21, e23) = op2(e21, e24)) & ~ (op2(e21, e22) = op2(e21, e24)) & ~ (op2(e21, e22) = op2(e21, e23)) & ~ (op2(e21, e21) = op2(e21, e24)) & ~ (op2(e21, e21) = op2(e21, e23)) & ~ (op2(e21, e21) = op2(e21, e22)) & ~ (op2(e21, e20) = op2(e21, e24)) & ~ (op2(e21, e20) = op2(e21, e23)) & ~ (op2(e21, e20) = op2(e21, e22)) & ~ (op2(e21, e20) = op2(e21, e21)) & ~ (op2(e20, e23) = op2(e20, e24)) & ~ (op2(e20, e22) = op2(e20, e24)) & ~ (op2(e20, e22) = op2(e20, e23)) & ~ (op2(e20, e21) = op2(e20, e24)) & ~ (op2(e20, e21) = op2(e20, e23)) & ~ (op2(e20, e21) = op2(e20, e22)) & ~ (op2(e20, e20) = op2(e20, e24)) & ~ (op2(e20, e20) = op2(e20, e23)) & ~ (op2(e20, e20) = op2(e20, e22)) & ~ (op2(e20, e20) = op2(e20, e21)) & ~ (op2(e23, e24) = op2(e24, e24)) & ~ (op2(e22, e24) = op2(e24, e24)) & ~ (op2(e22, e24) = op2(e23, e24)) & ~ (op2(e21, e24) = op2(e24, e24)) & ~ (op2(e21, e24) = op2(e23, e24)) & ~ (op2(e21, e24) = op2(e22, e24)) & ~ (op2(e20, e24) = op2(e24, e24)) & ~ (op2(e20, e24) = op2(e23, e24)) & ~ (op2(e20, e24) = op2(e22, e24)) & ~ (op2(e20, e24) = op2(e21, e24)) & ~ (op2(e23, e23) = op2(e24, e23)) & ~ (op2(e22, e23) = op2(e24, e23)) & ~ (op2(e22, e23) = op2(e23, e23)) & ~ (op2(e21, e23) = op2(e24, e23)) & ~ (op2(e21, e23) = op2(e23, e23)) & ~ (op2(e21, e23) = op2(e22, e23)) & ~ (op2(e20, e23) = op2(e24, e23)) & ~ (op2(e20, e23) = op2(e23, e23)) & ~ (op2(e20, e23) = op2(e22, e23)) & ~ (op2(e20, e23) = op2(e21, e23)) & ~ (op2(e23, e22) = op2(e24, e22)) & ~ (op2(e22, e22) = op2(e24, e22)) & ~ (op2(e22, e22) = op2(e23, e22)) & ~ (op2(e21, e22) = op2(e24, e22)) & ~ (op2(e21, e22) = op2(e23, e22)) & ~ (op2(e21, e22) = op2(e22, e22)) & ~ (op2(e20, e22) = op2(e24, e22)) & ~ (op2(e20, e22) = op2(e23, e22)) & ~ (op2(e20, e22) = op2(e22, e22)) & ~ (op2(e20, e22) = op2(e21, e22)) & ~ (op2(e23, e21) = op2(e24, e21)) & ~ (op2(e22, e21) = op2(e24, e21)) & ~ (op2(e22, e21) = op2(e23, e21)) & ~ (op2(e21, e21) = op2(e24, e21)) & ~ (op2(e21, e21) = op2(e23, e21)) & ~ (op2(e21, e21) = op2(e22, e21)) & ~ (op2(e20, e21) = op2(e24, e21)) & ~ (op2(e20, e21) = op2(e23, e21)) & ~ (op2(e20, e21) = op2(e22, e21)) & ~ (op2(e20, e21) = op2(e21, e21)) & ~ (op2(e23, e20) = op2(e24, e20)) & ~ (op2(e22, e20) = op2(e24, e20)) & ~ (op2(e22, e20) = op2(e23, e20)) & ~ (op2(e21, e20) = op2(e24, e20)) & ~ (op2(e21, e20) = op2(e23, e20)) & ~ (op2(e21, e20) = op2(e22, e20)) & ~ (op2(e20, e20) = op2(e24, e20)) & ~ (op2(e20, e20) = op2(e23, e20)) & ~ (op2(e20, e20) = op2(e22, e20)) & ~ (op2(e20, e20) = op2(e21, e20))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG049+1.p', ax8)).
fof(f11638, plain, (~ spl80_121 | ~ spl80_251 | ~ spl80_296 | spl80_521), inference(avatar_contradiction_clause, [], [f11637])).
fof(f11637, plain, ($false | (~ spl80_121 | ~ spl80_251 | ~ spl80_296 | spl80_521)), inference(subsumption_resolution, [], [f11636, f2376])).
fof(f2376, plain, ((e20 = op2(e20, e20)) | ~ spl80_251), inference(avatar_component_clause, [], [f2374])).
fof(f2374, plain, (spl80_251 <=> (e20 = op2(e20, e20))), introduced(avatar_definition, [new_symbols(naming, [spl80_251])])).
fof(f11636, plain, (~ (e20 = op2(e20, e20)) | (~ spl80_121 | ~ spl80_296 | spl80_521)), inference(backward_demodulation, [], [f8721, f2858])).
fof(f8721, plain, (~ (h18(e10) = op2(h18(e10), h18(e10))) | (~ spl80_121 | spl80_521)), inference(backward_demodulation, [], [f4497, f1780])).
fof(f1780, plain, ((e10 = op1(e10, e10)) | ~ spl80_121), inference(avatar_component_clause, [], [f1778])).
fof(f1778, plain, (spl80_121 <=> (e10 = op1(e10, e10))), introduced(avatar_definition, [new_symbols(naming, [spl80_121])])).
fof(f4497, plain, (~ (h18(op1(e10, e10)) = op2(h18(e10), h18(e10))) | spl80_521), inference(avatar_component_clause, [], [f4495])).
fof(f4495, plain, (spl80_521 <=> (h18(op1(e10, e10)) = op2(h18(e10), h18(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl80_521])])).
fof(f11549, plain, (spl80_296 | ~ spl80_191 | ~ spl80_291), inference(avatar_split_clause, [], [f11548, f2828, f2122, f2857])).
fof(f2122, plain, (spl80_191 <=> (e20 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl80_191])])).
fof(f11548, plain, ((e20 = h18(e10)) | (~ spl80_191 | ~ spl80_291)), inference(forward_demodulation, [], [f11545, f2124])).
fof(f2124, plain, ((e20 = op2(e22, e22)) | ~ spl80_191), inference(avatar_component_clause, [], [f2122])).
fof(f11545, plain, ((op2(e22, e22) = h18(e10)) | ~ spl80_291), inference(backward_demodulation, [], [f2644, f2829])).
fof(f2644, plain, (h18(e10) = op2(h8(e13), h8(e13))), inference(forward_demodulation, [], [f760, f712])).
fof(f760, plain, (h18(e10) = op2(op2(e21, e24), op2(e21, e24))), inference(cnf_transformation, [], [f33])).
fof(f11522, plain, (spl80_402 | ~ spl80_286 | ~ spl80_344 | ~ spl80_448), inference(avatar_split_clause, [], [f11519, f3778, f3141, f2798, f3499])).
fof(f3778, plain, (spl80_448 <=> (e20 = h3(e10))), introduced(avatar_definition, [new_symbols(naming, [spl80_448])])).
fof(f11519, plain, ((e20 = h8(e10)) | (~ spl80_286 | ~ spl80_344 | ~ spl80_448)), inference(backward_demodulation, [], [f11355, f3779])).
fof(f3779, plain, ((e20 = h3(e10)) | ~ spl80_448), inference(avatar_component_clause, [], [f3778])).
fof(f11355, plain, ((h3(e10) = h8(e10)) | (~ spl80_286 | ~ spl80_344)), inference(forward_demodulation, [], [f11350, f11325])).
fof(f11325, plain, ((op2(e23, e23) = h3(e10)) | ~ spl80_344), inference(backward_demodulation, [], [f2553, f3142])).
fof(f2553, plain, (h3(e10) = op2(h3(e11), h3(e11))), inference(forward_demodulation, [], [f685, f686])).
fof(f685, plain, (h3(e10) = op2(op2(e23, e20), op2(e23, e20))), inference(cnf_transformation, [], [f18])).
fof(f11350, plain, ((op2(e23, e23) = h8(e10)) | ~ spl80_286), inference(backward_demodulation, [], [f2600, f2799])).
fof(f2600, plain, (h8(e10) = op2(h8(e11), h8(e11))), inference(forward_demodulation, [], [f710, f711])).
fof(f710, plain, (h8(e10) = op2(op2(e24, e21), op2(e24, e21))), inference(cnf_transformation, [], [f23])).
fof(f11493, plain, (spl80_371 | ~ spl80_219), inference(avatar_split_clause, [], [f11400, f2239, f3314])).
fof(f3314, plain, (spl80_371 <=> (e23 = h6(e13))), introduced(avatar_definition, [new_symbols(naming, [spl80_371])])).
fof(f11400, plain, ((e23 = h6(e13)) | ~ spl80_219), inference(backward_demodulation, [], [f702, f2241])).
fof(f702, plain, (op2(e21, e22) = h6(e13)), inference(cnf_transformation, [], [f21])).
fof(f11468, plain, (spl80_291 | ~ spl80_208), inference(avatar_split_clause, [], [f11407, f2193, f2828])).
fof(f2193, plain, (spl80_208 <=> (e22 = op2(e21, e24))), introduced(avatar_definition, [new_symbols(naming, [spl80_208])])).
fof(f11407, plain, ((e22 = h8(e13)) | ~ spl80_208), inference(backward_demodulation, [], [f712, f2195])).
fof(f2195, plain, ((e22 = op2(e21, e24)) | ~ spl80_208), inference(avatar_component_clause, [], [f2193])).
fof(f11467, plain, (~ spl80_214 | ~ spl80_345), inference(avatar_split_clause, [], [f11233, f3149, f2218])).
fof(f2218, plain, (spl80_214 <=> (e23 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl80_214])])).
fof(f3149, plain, (spl80_345 <=> (e23 = h3(e13))), introduced(avatar_definition, [new_symbols(naming, [spl80_345])])).
fof(f11233, plain, (~ (e23 = op2(e21, e23)) | ~ spl80_345), inference(forward_demodulation, [], [f2541, f3150])).
fof(f3150, plain, ((e23 = h3(e13)) | ~ spl80_345), inference(avatar_component_clause, [], [f3149])).
fof(f2541, plain, ~ (op2(e21, e23) = h3(e13)), inference(backward_demodulation, [], [f502, f687])).
fof(f687, plain, (op2(e20, e23) = h3(e13)), inference(cnf_transformation, [], [f18])).
fof(f502, plain, ~ (op2(e20, e23) = op2(e21, e23)), inference(cnf_transformation, [], [f8])).
fof(f11429, plain, (spl80_448 | ~ spl80_161 | ~ spl80_344), inference(avatar_split_clause, [], [f11428, f3141, f1996, f3778])).
fof(f1996, plain, (spl80_161 <=> (e20 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl80_161])])).
fof(f11428, plain, ((e20 = h3(e10)) | (~ spl80_161 | ~ spl80_344)), inference(backward_demodulation, [], [f11325, f1998])).
fof(f1998, plain, ((e20 = op2(e23, e23)) | ~ spl80_161), inference(avatar_component_clause, [], [f1996])).
fof(f11390, plain, (spl80_155 | ~ spl80_256), inference(avatar_split_clause, [], [f11380, f2395, f1970])).
fof(f2395, plain, (spl80_256 <=> (e20 = unit2)), introduced(avatar_definition, [new_symbols(naming, [spl80_256])])).
fof(f11380, plain, ((e24 = op2(e24, e20)) | ~ spl80_256), inference(backward_demodulation, [], [f320, f2397])).
fof(f2397, plain, ((e20 = unit2) | ~ spl80_256), inference(avatar_component_clause, [], [f2395])).
fof(f320, plain, (e24 = op2(e24, unit2)), inference(cnf_transformation, [], [f5])).
fof(f5, plain, (((e24 = unit2) | (e23 = unit2) | (e22 = unit2) | (e21 = unit2) | (e20 = unit2)) & (e24 = op2(e24, unit2)) & (e24 = op2(unit2, e24)) & (e23 = op2(e23, unit2)) & (e23 = op2(unit2, e23)) & (e22 = op2(e22, unit2)) & (e22 = op2(unit2, e22)) & (e21 = op2(e21, unit2)) & (e21 = op2(unit2, e21)) & (e20 = op2(e20, unit2)) & (e20 = op2(unit2, e20))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG049+1.p', ax5)).
fof(f11389, plain, (spl80_235 | ~ spl80_256), inference(avatar_split_clause, [], [f11379, f2395, f2306])).
fof(f11379, plain, ((e24 = op2(e20, e24)) | ~ spl80_256), inference(backward_demodulation, [], [f319, f2397])).
fof(f319, plain, (e24 = op2(unit2, e24)), inference(cnf_transformation, [], [f5])).
fof(f11388, plain, (spl80_179 | ~ spl80_256), inference(avatar_split_clause, [], [f11378, f2395, f2071])).
fof(f11378, plain, ((e23 = op2(e23, e20)) | ~ spl80_256), inference(backward_demodulation, [], [f318, f2397])).
fof(f318, plain, (e23 = op2(e23, unit2)), inference(cnf_transformation, [], [f5])).
fof(f11387, plain, (spl80_247 | ~ spl80_256), inference(avatar_split_clause, [], [f11373, f2395, f2357])).
fof(f11373, plain, ((e21 = op2(e20, e21)) | ~ spl80_256), inference(backward_demodulation, [], [f313, f2397])).
fof(f313, plain, (e21 = op2(unit2, e21)), inference(cnf_transformation, [], [f5])).
fof(f11386, plain, (spl80_251 | ~ spl80_256), inference(avatar_split_clause, [], [f11372, f2395, f2374])).
fof(f11372, plain, ((e20 = op2(e20, e20)) | ~ spl80_256), inference(backward_demodulation, [], [f312, f2397])).
fof(f312, plain, (e20 = op2(e20, unit2)), inference(cnf_transformation, [], [f5])).
fof(f11313, plain, (~ spl80_232 | ~ spl80_429), inference(avatar_split_clause, [], [f11305, f3662, f2294])).
fof(f2294, plain, (spl80_232 <=> (e21 = op2(e20, e24))), introduced(avatar_definition, [new_symbols(naming, [spl80_232])])).
fof(f3662, plain, (spl80_429 <=> (e21 = h1(e13))), introduced(avatar_definition, [new_symbols(naming, [spl80_429])])).
fof(f11305, plain, (~ (e21 = op2(e20, e24)) | ~ spl80_429), inference(backward_demodulation, [], [f2516, f3663])).
fof(f3663, plain, ((e21 = h1(e13)) | ~ spl80_429), inference(avatar_component_clause, [], [f3662])).
fof(f2516, plain, ~ (op2(e20, e24) = h1(e13)), inference(backward_demodulation, [], [f528, f677])).
fof(f677, plain, (op2(e20, e21) = h1(e13)), inference(cnf_transformation, [], [f16])).
fof(f16, plain, ((op2(e20, e21) = h1(e13)) & (op2(e21, e20) = h1(e11)) & (h1(e10) = op2(op2(e21, e20), op2(e21, e20))) & (e21 = h1(e14)) & (e20 = h1(e12))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG049+1.p', ax16)).
fof(f528, plain, ~ (op2(e20, e21) = op2(e20, e24)), inference(cnf_transformation, [], [f8])).
fof(f11199, plain, (~ spl80_218 | ~ spl80_387), inference(avatar_split_clause, [], [f11198, f3409, f2235])).
fof(f2235, plain, (spl80_218 <=> (e22 = op2(e21, e22))), introduced(avatar_definition, [new_symbols(naming, [spl80_218])])).
fof(f3409, plain, (spl80_387 <=> (e22 = h2(e13))), introduced(avatar_definition, [new_symbols(naming, [spl80_387])])).
fof(f11198, plain, (~ (e22 = op2(e21, e22)) | ~ spl80_387), inference(forward_demodulation, [], [f2526, f3410])).
fof(f3410, plain, ((e22 = h2(e13)) | ~ spl80_387), inference(avatar_component_clause, [], [f3409])).
fof(f2526, plain, ~ (op2(e21, e22) = h2(e13)), inference(backward_demodulation, [], [f492, f682])).
fof(f682, plain, (op2(e20, e22) = h2(e13)), inference(cnf_transformation, [], [f17])).
fof(f17, plain, ((op2(e20, e22) = h2(e13)) & (op2(e22, e20) = h2(e11)) & (h2(e10) = op2(op2(e22, e20), op2(e22, e20))) & (e22 = h2(e14)) & (e20 = h2(e12))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG049+1.p', ax17)).
fof(f492, plain, ~ (op2(e20, e22) = op2(e21, e22)), inference(cnf_transformation, [], [f8])).
fof(f11193, plain, (~ spl80_198 | ~ spl80_386), inference(avatar_split_clause, [], [f11192, f3401, f2151])).
fof(f2151, plain, (spl80_198 <=> (e22 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl80_198])])).
fof(f3401, plain, (spl80_386 <=> (e22 = h2(e11))), introduced(avatar_definition, [new_symbols(naming, [spl80_386])])).
fof(f11192, plain, (~ (e22 = op2(e22, e21)) | ~ spl80_386), inference(forward_demodulation, [], [f2536, f3402])).
fof(f3402, plain, ((e22 = h2(e11)) | ~ spl80_386), inference(avatar_component_clause, [], [f3401])).
fof(f2536, plain, ~ (op2(e22, e21) = h2(e11)), inference(backward_demodulation, [], [f542, f681])).
fof(f681, plain, (op2(e22, e20) = h2(e11)), inference(cnf_transformation, [], [f17])).
fof(f542, plain, ~ (op2(e22, e20) = op2(e22, e21)), inference(cnf_transformation, [], [f8])).
fof(f11186, plain, (~ spl80_187 | ~ spl80_188), inference(avatar_contradiction_clause, [], [f11185])).
fof(f11185, plain, ($false | (~ spl80_187 | ~ spl80_188)), inference(subsumption_resolution, [], [f11184, f586])).
fof(f586, plain, ~ (e21 = e22), inference(cnf_transformation, [], [f10])).
fof(f11184, plain, ((e21 = e22) | (~ spl80_187 | ~ spl80_188)), inference(forward_demodulation, [], [f2111, f2107])).
fof(f2111, plain, ((e22 = op2(e22, e23)) | ~ spl80_188), inference(avatar_component_clause, [], [f2109])).
fof(f2109, plain, (spl80_188 <=> (e22 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl80_188])])).
fof(f11060, plain, (spl80_25 | ~ spl80_126), inference(avatar_split_clause, [], [f11045, f1799, f1374])).
fof(f1799, plain, (spl80_126 <=> (e10 = unit1)), introduced(avatar_definition, [new_symbols(naming, [spl80_126])])).
fof(f11045, plain, ((e14 = op1(e14, e10)) | ~ spl80_126), inference(backward_demodulation, [], [f234, f1801])).
fof(f1801, plain, ((e10 = unit1) | ~ spl80_126), inference(avatar_component_clause, [], [f1799])).
fof(f234, plain, (e14 = op1(e14, unit1)), inference(cnf_transformation, [], [f2])).
fof(f2, plain, (((e14 = unit1) | (e13 = unit1) | (e12 = unit1) | (e11 = unit1) | (e10 = unit1)) & (e14 = op1(e14, unit1)) & (e14 = op1(unit1, e14)) & (e13 = op1(e13, unit1)) & (e13 = op1(unit1, e13)) & (e12 = op1(e12, unit1)) & (e12 = op1(unit1, e12)) & (e11 = op1(e11, unit1)) & (e11 = op1(unit1, e11)) & (e10 = op1(e10, unit1)) & (e10 = op1(unit1, e10))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG049+1.p', ax2)).
fof(f11059, plain, (~ spl80_102 | ~ spl80_126), inference(avatar_contradiction_clause, [], [f11058])).
fof(f11058, plain, ($false | (~ spl80_102 | ~ spl80_126)), inference(subsumption_resolution, [], [f11057, f578])).
fof(f11057, plain, ((e11 = e14) | (~ spl80_102 | ~ spl80_126)), inference(forward_demodulation, [], [f11044, f1700])).
fof(f1700, plain, ((e11 = op1(e10, e14)) | ~ spl80_102), inference(avatar_component_clause, [], [f1698])).
fof(f1698, plain, (spl80_102 <=> (e11 = op1(e10, e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_102])])).
fof(f11044, plain, ((e14 = op1(e10, e14)) | ~ spl80_126), inference(backward_demodulation, [], [f233, f1801])).
fof(f233, plain, (e14 = op1(unit1, e14)), inference(cnf_transformation, [], [f2])).
fof(f11056, plain, (~ spl80_48 | ~ spl80_126), inference(avatar_contradiction_clause, [], [f11055])).
fof(f11055, plain, ($false | (~ spl80_48 | ~ spl80_126)), inference(subsumption_resolution, [], [f11054, f579])).
fof(f579, plain, ~ (e12 = e13), inference(cnf_transformation, [], [f9])).
fof(f11054, plain, ((e12 = e13) | (~ spl80_48 | ~ spl80_126)), inference(forward_demodulation, [], [f11043, f1473])).
fof(f1473, plain, ((e12 = op1(e13, e10)) | ~ spl80_48), inference(avatar_component_clause, [], [f1471])).
fof(f1471, plain, (spl80_48 <=> (e12 = op1(e13, e10))), introduced(avatar_definition, [new_symbols(naming, [spl80_48])])).
fof(f11043, plain, ((e13 = op1(e13, e10)) | ~ spl80_126), inference(backward_demodulation, [], [f232, f1801])).
fof(f232, plain, (e13 = op1(e13, unit1)), inference(cnf_transformation, [], [f2])).
fof(f11053, plain, (spl80_109 | ~ spl80_126), inference(avatar_split_clause, [], [f11042, f1799, f1727])).
fof(f11042, plain, ((e13 = op1(e10, e13)) | ~ spl80_126), inference(backward_demodulation, [], [f231, f1801])).
fof(f231, plain, (e13 = op1(unit1, e13)), inference(cnf_transformation, [], [f2])).
fof(f11052, plain, (spl80_73 | ~ spl80_126), inference(avatar_split_clause, [], [f11041, f1799, f1576])).
fof(f11041, plain, ((e12 = op1(e12, e10)) | ~ spl80_126), inference(backward_demodulation, [], [f230, f1801])).
fof(f230, plain, (e12 = op1(e12, unit1)), inference(cnf_transformation, [], [f2])).
fof(f11051, plain, (spl80_113 | ~ spl80_126), inference(avatar_split_clause, [], [f11040, f1799, f1744])).
fof(f11040, plain, ((e12 = op1(e10, e12)) | ~ spl80_126), inference(backward_demodulation, [], [f229, f1801])).
fof(f229, plain, (e12 = op1(unit1, e12)), inference(cnf_transformation, [], [f2])).
fof(f11050, plain, (spl80_117 | ~ spl80_126), inference(avatar_split_clause, [], [f11038, f1799, f1761])).
fof(f11038, plain, ((e11 = op1(e10, e11)) | ~ spl80_126), inference(backward_demodulation, [], [f227, f1801])).
fof(f227, plain, (e11 = op1(unit1, e11)), inference(cnf_transformation, [], [f2])).
fof(f11007, plain, (~ spl80_186 | ~ spl80_187), inference(avatar_contradiction_clause, [], [f11006])).
fof(f11006, plain, ($false | (~ spl80_186 | ~ spl80_187)), inference(subsumption_resolution, [], [f11005, f582])).
fof(f582, plain, ~ (e20 = e21), inference(cnf_transformation, [], [f10])).
fof(f11005, plain, ((e20 = e21) | (~ spl80_186 | ~ spl80_187)), inference(forward_demodulation, [], [f2107, f2103])).
fof(f2103, plain, ((e20 = op2(e22, e23)) | ~ spl80_186), inference(avatar_component_clause, [], [f2101])).
fof(f2101, plain, (spl80_186 <=> (e20 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl80_186])])).
fof(f10962, plain, (~ spl80_221 | ~ spl80_257), inference(avatar_contradiction_clause, [], [f10961])).
fof(f10961, plain, ($false | (~ spl80_221 | ~ spl80_257)), inference(subsumption_resolution, [], [f10960, f582])).
fof(f10960, plain, ((e20 = e21) | (~ spl80_221 | ~ spl80_257)), inference(forward_demodulation, [], [f10945, f2250])).
fof(f10945, plain, ((e21 = op2(e21, e21)) | ~ spl80_257), inference(backward_demodulation, [], [f314, f2401])).
fof(f2401, plain, ((e21 = unit2) | ~ spl80_257), inference(avatar_component_clause, [], [f2399])).
fof(f2399, plain, (spl80_257 <=> (e21 = unit2)), introduced(avatar_definition, [new_symbols(naming, [spl80_257])])).
fof(f314, plain, (e21 = op2(e21, unit2)), inference(cnf_transformation, [], [f5])).
fof(f10778, plain, (~ spl80_295 | ~ spl80_428 | ~ spl80_464), inference(avatar_split_clause, [], [f10777, f3876, f3654, f2852])).
fof(f2852, plain, (spl80_295 <=> (e20 = h8(e13))), introduced(avatar_definition, [new_symbols(naming, [spl80_295])])).
fof(f3654, plain, (spl80_428 <=> (e21 = h1(e11))), introduced(avatar_definition, [new_symbols(naming, [spl80_428])])).
fof(f3876, plain, (spl80_464 <=> (e20 = h1(e10))), introduced(avatar_definition, [new_symbols(naming, [spl80_464])])).
fof(f10777, plain, (~ (e20 = h8(e13)) | (~ spl80_428 | ~ spl80_464)), inference(forward_demodulation, [], [f10064, f3877])).
fof(f3877, plain, ((e20 = h1(e10)) | ~ spl80_464), inference(avatar_component_clause, [], [f3876])).
fof(f10064, plain, (~ (h1(e10) = h8(e13)) | ~ spl80_428), inference(backward_demodulation, [], [f2595, f10054])).
fof(f10054, plain, ((op2(e21, e21) = h1(e10)) | ~ spl80_428), inference(backward_demodulation, [], [f2525, f3655])).
fof(f3655, plain, ((e21 = h1(e11)) | ~ spl80_428), inference(avatar_component_clause, [], [f3654])).
fof(f2525, plain, (h1(e10) = op2(h1(e11), h1(e11))), inference(forward_demodulation, [], [f675, f676])).
fof(f676, plain, (op2(e21, e20) = h1(e11)), inference(cnf_transformation, [], [f16])).
fof(f675, plain, (h1(e10) = op2(op2(e21, e20), op2(e21, e20))), inference(cnf_transformation, [], [f16])).
fof(f2595, plain, ~ (op2(e21, e21) = h8(e13)), inference(backward_demodulation, [], [f538, f712])).
fof(f538, plain, ~ (op2(e21, e21) = op2(e21, e24)), inference(cnf_transformation, [], [f8])).
fof(f10707, plain, (~ spl80_83 | ~ spl80_78), inference(avatar_split_clause, [], [f10706, f1597, f1618])).
fof(f1618, plain, (spl80_83 <=> (e12 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl80_83])])).
fof(f10706, plain, (~ (e12 = op1(e11, e13)) | ~ spl80_78), inference(forward_demodulation, [], [f441, f1599])).
fof(f441, plain, ~ (op1(e11, e13) = op1(e11, e14)), inference(cnf_transformation, [], [f7])).
fof(f10699, plain, (~ spl80_45 | ~ spl80_40), inference(avatar_split_clause, [], [f10698, f1437, f1458])).
fof(f1458, plain, (spl80_45 <=> (e14 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl80_45])])).
fof(f10698, plain, (~ (e14 = op1(e13, e11)) | ~ spl80_40), inference(forward_demodulation, [], [f456, f1439])).
fof(f456, plain, ~ (op1(e13, e11) = op1(e13, e12)), inference(cnf_transformation, [], [f7])).
fof(f10693, plain, (~ spl80_41 | ~ spl80_91), inference(avatar_split_clause, [], [f9828, f1652, f1442])).
fof(f1442, plain, (spl80_41 <=> (e10 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl80_41])])).
fof(f9828, plain, (~ (e10 = op1(e13, e11)) | ~ spl80_91), inference(backward_demodulation, [], [f387, f1654])).
fof(f387, plain, ~ (op1(e11, e11) = op1(e13, e11)), inference(cnf_transformation, [], [f7])).
fof(f10663, plain, (~ spl80_36 | ~ spl80_40), inference(avatar_contradiction_clause, [], [f10662])).
fof(f10662, plain, ($false | (~ spl80_36 | ~ spl80_40)), inference(subsumption_resolution, [], [f10661, f575])).
fof(f575, plain, ~ (e10 = e14), inference(cnf_transformation, [], [f9])).
fof(f10661, plain, ((e10 = e14) | (~ spl80_36 | ~ spl80_40)), inference(forward_demodulation, [], [f1439, f1423])).
fof(f1423, plain, ((e10 = op1(e13, e12)) | ~ spl80_36), inference(avatar_component_clause, [], [f1421])).
fof(f1421, plain, (spl80_36 <=> (e10 = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl80_36])])).
fof(f10654, plain, (~ spl80_51 | ~ spl80_54), inference(avatar_contradiction_clause, [], [f10653])).
fof(f10653, plain, ($false | (~ spl80_51 | ~ spl80_54)), inference(subsumption_resolution, [], [f10651, f574])).
fof(f574, plain, ~ (e10 = e13), inference(cnf_transformation, [], [f9])).
fof(f10651, plain, ((e10 = e13) | (~ spl80_51 | ~ spl80_54)), inference(backward_demodulation, [], [f1498, f1486])).
fof(f1486, plain, ((e10 = op1(e12, e14)) | ~ spl80_51), inference(avatar_component_clause, [], [f1484])).
fof(f1484, plain, (spl80_51 <=> (e10 = op1(e12, e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_51])])).
fof(f10641, plain, (~ spl80_97 | ~ spl80_99), inference(avatar_contradiction_clause, [], [f10640])).
fof(f10640, plain, ($false | (~ spl80_97 | ~ spl80_99)), inference(subsumption_resolution, [], [f10639, f577])).
fof(f10639, plain, ((e11 = e13) | (~ spl80_97 | ~ spl80_99)), inference(forward_demodulation, [], [f1687, f1679])).
fof(f1687, plain, ((e13 = op1(e11, e10)) | ~ spl80_99), inference(avatar_component_clause, [], [f1685])).
fof(f1685, plain, (spl80_99 <=> (e13 = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl80_99])])).
fof(f10629, plain, (~ spl80_12 | ~ spl80_130), inference(avatar_contradiction_clause, [], [f10628])).
fof(f10628, plain, ($false | (~ spl80_12 | ~ spl80_130)), inference(subsumption_resolution, [], [f10627, f576])).
fof(f576, plain, ~ (e11 = e12), inference(cnf_transformation, [], [f9])).
fof(f10627, plain, ((e11 = e12) | (~ spl80_12 | ~ spl80_130)), inference(forward_demodulation, [], [f10614, f1322])).
fof(f10614, plain, ((e12 = op1(e14, e12)) | ~ spl80_130), inference(backward_demodulation, [], [f229, f1817])).
fof(f1817, plain, ((e14 = unit1) | ~ spl80_130), inference(avatar_component_clause, [], [f1815])).
fof(f1815, plain, (spl80_130 <=> (e14 = unit1)), introduced(avatar_definition, [new_symbols(naming, [spl80_130])])).
fof(f10555, plain, (~ spl80_227 | ~ spl80_229), inference(avatar_contradiction_clause, [], [f10554])).
fof(f10554, plain, ($false | (~ spl80_227 | ~ spl80_229)), inference(subsumption_resolution, [], [f10553, f587])).
fof(f587, plain, ~ (e21 = e23), inference(cnf_transformation, [], [f10])).
fof(f10553, plain, ((e21 = e23) | (~ spl80_227 | ~ spl80_229)), inference(forward_demodulation, [], [f2283, f2275])).
fof(f2283, plain, ((e23 = op2(e21, e20)) | ~ spl80_229), inference(avatar_component_clause, [], [f2281])).
fof(f2281, plain, (spl80_229 <=> (e23 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl80_229])])).
fof(f10548, plain, (~ spl80_238 | ~ spl80_239), inference(avatar_contradiction_clause, [], [f10547])).
fof(f10547, plain, ($false | (~ spl80_238 | ~ spl80_239)), inference(subsumption_resolution, [], [f10546, f589])).
fof(f589, plain, ~ (e22 = e23), inference(cnf_transformation, [], [f10])).
fof(f10546, plain, ((e22 = e23) | (~ spl80_238 | ~ spl80_239)), inference(backward_demodulation, [], [f2325, f2321])).
fof(f2321, plain, ((e22 = op2(e20, e23)) | ~ spl80_238), inference(avatar_component_clause, [], [f2319])).
fof(f2319, plain, (spl80_238 <=> (e22 = op2(e20, e23))), introduced(avatar_definition, [new_symbols(naming, [spl80_238])])).
fof(f10540, plain, (spl80_421 | ~ spl80_249), inference(avatar_split_clause, [], [f10539, f2365, f3615])).
fof(f3615, plain, (spl80_421 <=> (e23 = h1(e13))), introduced(avatar_definition, [new_symbols(naming, [spl80_421])])).
fof(f2365, plain, (spl80_249 <=> (e23 = op2(e20, e21))), introduced(avatar_definition, [new_symbols(naming, [spl80_249])])).
fof(f10539, plain, ((e23 = h1(e13)) | ~ spl80_249), inference(backward_demodulation, [], [f677, f2367])).
fof(f2367, plain, ((e23 = op2(e20, e21)) | ~ spl80_249), inference(avatar_component_clause, [], [f2365])).
fof(f10524, plain, (~ spl80_142 | ~ spl80_260), inference(avatar_contradiction_clause, [], [f10523])).
fof(f10523, plain, ($false | (~ spl80_142 | ~ spl80_260)), inference(subsumption_resolution, [], [f10522, f586])).
fof(f10522, plain, ((e21 = e22) | (~ spl80_142 | ~ spl80_260)), inference(forward_demodulation, [], [f10515, f1918])).
fof(f10515, plain, ((e22 = op2(e24, e22)) | ~ spl80_260), inference(backward_demodulation, [], [f315, f2413])).
fof(f2413, plain, ((e24 = unit2) | ~ spl80_260), inference(avatar_component_clause, [], [f2411])).
fof(f2411, plain, (spl80_260 <=> (e24 = unit2)), introduced(avatar_definition, [new_symbols(naming, [spl80_260])])).
fof(f315, plain, (e22 = op2(unit2, e22)), inference(cnf_transformation, [], [f5])).
fof(f10431, plain, (~ spl80_341 | ~ spl80_428 | ~ spl80_464), inference(avatar_split_clause, [], [f10427, f3876, f3654, f3125])).
fof(f3125, plain, (spl80_341 <=> (e20 = h7(e13))), introduced(avatar_definition, [new_symbols(naming, [spl80_341])])).
fof(f10427, plain, (~ (e20 = h7(e13)) | (~ spl80_428 | ~ spl80_464)), inference(backward_demodulation, [], [f10062, f3877])).
fof(f10062, plain, (~ (h1(e10) = h7(e13)) | ~ spl80_428), inference(backward_demodulation, [], [f2584, f10054])).
fof(f2584, plain, ~ (op2(e21, e21) = h7(e13)), inference(backward_demodulation, [], [f537, f707])).
fof(f707, plain, (op2(e21, e23) = h7(e13)), inference(cnf_transformation, [], [f22])).
fof(f22, plain, ((op2(e21, e23) = h7(e13)) & (op2(e23, e21) = h7(e11)) & (h7(e10) = op2(op2(e23, e21), op2(e23, e21))) & (e23 = h7(e14)) & (e21 = h7(e12))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG049+1.p', ax22)).
fof(f537, plain, ~ (op2(e21, e21) = op2(e21, e23)), inference(cnf_transformation, [], [f8])).
fof(f10430, plain, (~ spl80_378 | ~ spl80_428 | ~ spl80_464), inference(avatar_split_clause, [], [f10426, f3876, f3654, f3355])).
fof(f3355, plain, (spl80_378 <=> (e20 = h6(e11))), introduced(avatar_definition, [new_symbols(naming, [spl80_378])])).
fof(f10426, plain, (~ (e20 = h6(e11)) | (~ spl80_428 | ~ spl80_464)), inference(backward_demodulation, [], [f10061, f3877])).
fof(f10061, plain, (~ (h1(e10) = h6(e11)) | ~ spl80_428), inference(backward_demodulation, [], [f2574, f10054])).
fof(f2574, plain, ~ (op2(e21, e21) = h6(e11)), inference(backward_demodulation, [], [f486, f701])).
fof(f486, plain, ~ (op2(e21, e21) = op2(e22, e21)), inference(cnf_transformation, [], [f8])).
fof(f10381, plain, (spl80_438 | ~ spl80_231), inference(avatar_split_clause, [], [f10102, f2290, f3717])).
fof(f3717, plain, (spl80_438 <=> (e20 = h4(e13))), introduced(avatar_definition, [new_symbols(naming, [spl80_438])])).
fof(f2290, plain, (spl80_231 <=> (e20 = op2(e20, e24))), introduced(avatar_definition, [new_symbols(naming, [spl80_231])])).
fof(f10102, plain, ((e20 = h4(e13)) | ~ spl80_231), inference(backward_demodulation, [], [f692, f2292])).
fof(f2292, plain, ((e20 = op2(e20, e24)) | ~ spl80_231), inference(avatar_component_clause, [], [f2290])).
fof(f692, plain, (op2(e20, e24) = h4(e13)), inference(cnf_transformation, [], [f19])).
fof(f19, plain, ((op2(e20, e24) = h4(e13)) & (op2(e24, e20) = h4(e11)) & (h4(e10) = op2(op2(e24, e20), op2(e24, e20))) & (e24 = h4(e14)) & (e20 = h4(e12))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG049+1.p', ax19)).
fof(f10376, plain, (~ spl80_239 | spl80_345), inference(avatar_contradiction_clause, [], [f10375])).
fof(f10375, plain, ($false | (~ spl80_239 | spl80_345)), inference(subsumption_resolution, [], [f10374, f3151])).
fof(f3151, plain, (~ (e23 = h3(e13)) | spl80_345), inference(avatar_component_clause, [], [f3149])).
fof(f10374, plain, ((e23 = h3(e13)) | ~ spl80_239), inference(forward_demodulation, [], [f687, f2325])).
fof(f10373, plain, (~ spl80_421 | ~ spl80_239), inference(avatar_split_clause, [], [f10372, f2323, f3615])).
fof(f10372, plain, (~ (e23 = h1(e13)) | ~ spl80_239), inference(forward_demodulation, [], [f2515, f2325])).
fof(f2515, plain, ~ (op2(e20, e23) = h1(e13)), inference(backward_demodulation, [], [f527, f677])).
fof(f527, plain, ~ (op2(e20, e21) = op2(e20, e23)), inference(cnf_transformation, [], [f8])).
fof(f10369, plain, (~ spl80_212 | ~ spl80_428), inference(avatar_split_clause, [], [f10368, f3654, f2210])).
fof(f2210, plain, (spl80_212 <=> (e21 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl80_212])])).
fof(f10368, plain, (~ (e21 = op2(e21, e23)) | ~ spl80_428), inference(forward_demodulation, [], [f2523, f3655])).
fof(f2523, plain, ~ (op2(e21, e23) = h1(e11)), inference(backward_demodulation, [], [f534, f676])).
fof(f534, plain, ~ (op2(e21, e20) = op2(e21, e23)), inference(cnf_transformation, [], [f8])).
fof(f10352, plain, (~ spl80_179 | spl80_344), inference(avatar_contradiction_clause, [], [f10351])).
fof(f10351, plain, ($false | (~ spl80_179 | spl80_344)), inference(subsumption_resolution, [], [f10350, f3143])).
fof(f3143, plain, (~ (e23 = h3(e11)) | spl80_344), inference(avatar_component_clause, [], [f3141])).
fof(f10350, plain, ((e23 = h3(e11)) | ~ spl80_179), inference(forward_demodulation, [], [f686, f2073])).
fof(f10336, plain, (~ spl80_132 | ~ spl80_278), inference(avatar_split_clause, [], [f9639, f2752, f1874])).
fof(f1874, plain, (spl80_132 <=> (e21 = op2(e24, e24))), introduced(avatar_definition, [new_symbols(naming, [spl80_132])])).
fof(f2752, plain, (spl80_278 <=> (e21 = h12(e11))), introduced(avatar_definition, [new_symbols(naming, [spl80_278])])).
fof(f9639, plain, (~ (e21 = op2(e24, e24)) | ~ spl80_278), inference(backward_demodulation, [], [f2621, f2753])).
fof(f2753, plain, ((e21 = h12(e11)) | ~ spl80_278), inference(avatar_component_clause, [], [f2752])).
fof(f2621, plain, ~ (op2(e24, e24) = h12(e11)), inference(backward_demodulation, [], [f570, f731])).
fof(f731, plain, (op2(e24, e22) = h12(e11)), inference(cnf_transformation, [], [f27])).
fof(f27, plain, ((op2(e22, e24) = h12(e13)) & (op2(e24, e22) = h12(e11)) & (op2(op2(e24, e22), op2(e24, e22)) = h12(e10)) & (e24 = h12(e14)) & (e22 = h12(e12))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG049+1.p', ax27)).
fof(f570, plain, ~ (op2(e24, e22) = op2(e24, e24)), inference(cnf_transformation, [], [f8])).
fof(f10242, plain, (spl80_82 | ~ spl80_129), inference(avatar_contradiction_clause, [], [f10241])).
fof(f10241, plain, ($false | (spl80_82 | ~ spl80_129)), inference(subsumption_resolution, [], [f10230, f1615])).
fof(f1615, plain, (~ (e11 = op1(e11, e13)) | spl80_82), inference(avatar_component_clause, [], [f1614])).
fof(f1614, plain, (spl80_82 <=> (e11 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl80_82])])).
fof(f10230, plain, ((e11 = op1(e11, e13)) | ~ spl80_129), inference(backward_demodulation, [], [f228, f1813])).
fof(f1813, plain, ((e13 = unit1) | ~ spl80_129), inference(avatar_component_clause, [], [f1811])).
fof(f1811, plain, (spl80_129 <=> (e13 = unit1)), introduced(avatar_definition, [new_symbols(naming, [spl80_129])])).
fof(f228, plain, (e11 = op1(e11, unit1)), inference(cnf_transformation, [], [f2])).
fof(f10202, plain, (~ spl80_142 | ~ spl80_145), inference(avatar_contradiction_clause, [], [f10201])).
fof(f10201, plain, ($false | (~ spl80_142 | ~ spl80_145)), inference(subsumption_resolution, [], [f10200, f588])).
fof(f10200, plain, ((e21 = e24) | (~ spl80_142 | ~ spl80_145)), inference(forward_demodulation, [], [f1930, f1918])).
fof(f1930, plain, ((e24 = op2(e24, e22)) | ~ spl80_145), inference(avatar_component_clause, [], [f1928])).
fof(f1928, plain, (spl80_145 <=> (e24 = op2(e24, e22))), introduced(avatar_definition, [new_symbols(naming, [spl80_145])])).
fof(f10121, plain, (spl80_464 | ~ spl80_221 | ~ spl80_428), inference(avatar_split_clause, [], [f10120, f3654, f2248, f3876])).
fof(f10120, plain, ((e20 = h1(e10)) | (~ spl80_221 | ~ spl80_428)), inference(backward_demodulation, [], [f10054, f2250])).
fof(f10115, plain, (~ spl80_227 | ~ spl80_230), inference(avatar_contradiction_clause, [], [f10114])).
fof(f10114, plain, ($false | (~ spl80_227 | ~ spl80_230)), inference(subsumption_resolution, [], [f10113, f588])).
fof(f10113, plain, ((e21 = e24) | (~ spl80_227 | ~ spl80_230)), inference(backward_demodulation, [], [f2287, f2275])).
fof(f2287, plain, ((e24 = op2(e21, e20)) | ~ spl80_230), inference(avatar_component_clause, [], [f2285])).
fof(f2285, plain, (spl80_230 <=> (e24 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl80_230])])).
fof(f10100, plain, (~ spl80_237 | ~ spl80_239), inference(avatar_contradiction_clause, [], [f10099])).
fof(f10099, plain, ($false | (~ spl80_237 | ~ spl80_239)), inference(subsumption_resolution, [], [f10098, f587])).
fof(f10098, plain, ((e21 = e23) | (~ spl80_237 | ~ spl80_239)), inference(forward_demodulation, [], [f2325, f2317])).
fof(f2317, plain, ((e21 = op2(e20, e23)) | ~ spl80_237), inference(avatar_component_clause, [], [f2315])).
fof(f2315, plain, (spl80_237 <=> (e21 = op2(e20, e23))), introduced(avatar_definition, [new_symbols(naming, [spl80_237])])).
fof(f10097, plain, (~ spl80_425 | ~ spl80_243), inference(avatar_split_clause, [], [f10094, f2340, f3638])).
fof(f3638, plain, (spl80_425 <=> (e22 = h1(e13))), introduced(avatar_definition, [new_symbols(naming, [spl80_425])])).
fof(f10094, plain, (~ (e22 = h1(e13)) | ~ spl80_243), inference(backward_demodulation, [], [f2514, f2342])).
fof(f2514, plain, ~ (op2(e20, e22) = h1(e13)), inference(backward_demodulation, [], [f526, f677])).
fof(f526, plain, ~ (op2(e20, e21) = op2(e20, e22)), inference(cnf_transformation, [], [f8])).
fof(f10096, plain, (~ spl80_243 | spl80_387), inference(avatar_contradiction_clause, [], [f10095])).
fof(f10095, plain, ($false | (~ spl80_243 | spl80_387)), inference(subsumption_resolution, [], [f10093, f3411])).
fof(f3411, plain, (~ (e22 = h2(e13)) | spl80_387), inference(avatar_component_clause, [], [f3409])).
fof(f10093, plain, ((e22 = h2(e13)) | ~ spl80_243), inference(backward_demodulation, [], [f682, f2342])).
fof(f10081, plain, (spl80_212 | ~ spl80_259), inference(avatar_split_clause, [], [f10069, f2407, f2210])).
fof(f2407, plain, (spl80_259 <=> (e23 = unit2)), introduced(avatar_definition, [new_symbols(naming, [spl80_259])])).
fof(f10069, plain, ((e21 = op2(e21, e23)) | ~ spl80_259), inference(backward_demodulation, [], [f314, f2409])).
fof(f2409, plain, ((e23 = unit2) | ~ spl80_259), inference(avatar_component_clause, [], [f2407])).
fof(f10058, plain, (~ spl80_207 | ~ spl80_428), inference(avatar_split_clause, [], [f10053, f3654, f2189])).
fof(f2189, plain, (spl80_207 <=> (e21 = op2(e21, e24))), introduced(avatar_definition, [new_symbols(naming, [spl80_207])])).
fof(f10053, plain, (~ (e21 = op2(e21, e24)) | ~ spl80_428), inference(backward_demodulation, [], [f2524, f3655])).
fof(f2524, plain, ~ (op2(e21, e24) = h1(e11)), inference(backward_demodulation, [], [f535, f676])).
fof(f535, plain, ~ (op2(e21, e20) = op2(e21, e24)), inference(cnf_transformation, [], [f8])).
fof(f10030, plain, (spl80_221 | ~ spl80_278), inference(avatar_split_clause, [], [f9640, f2752, f2248])).
fof(f9640, plain, ((e20 = op2(e21, e21)) | ~ spl80_278), inference(backward_demodulation, [], [f2622, f2753])).
fof(f2622, plain, (e20 = op2(h12(e11), h12(e11))), inference(backward_demodulation, [], [f670, f731])).
fof(f670, plain, (e20 = op2(op2(e24, e22), op2(e24, e22))), inference(cnf_transformation, [], [f15])).
fof(f15, plain, ((e23 = op2(e22, e24)) & (e21 = op2(e24, e22)) & (e20 = op2(op2(e24, e22), op2(e24, e22)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG049+1.p', ax15)).
fof(f10019, plain, (~ spl80_203 | spl80_386), inference(avatar_contradiction_clause, [], [f10018])).
fof(f10018, plain, ($false | (~ spl80_203 | spl80_386)), inference(subsumption_resolution, [], [f10017, f3403])).
fof(f3403, plain, (~ (e22 = h2(e11)) | spl80_386), inference(avatar_component_clause, [], [f3401])).
fof(f10017, plain, ((e22 = h2(e11)) | ~ spl80_203), inference(forward_demodulation, [], [f681, f2174])).
fof(f10016, plain, (~ spl80_424 | ~ spl80_203), inference(avatar_split_clause, [], [f10015, f2172, f3631])).
fof(f3631, plain, (spl80_424 <=> (e22 = h1(e11))), introduced(avatar_definition, [new_symbols(naming, [spl80_424])])).
fof(f10015, plain, (~ (e22 = h1(e11)) | ~ spl80_203), inference(forward_demodulation, [], [f2518, f2174])).
fof(f2518, plain, ~ (op2(e22, e20) = h1(e11)), inference(backward_demodulation, [], [f476, f676])).
fof(f476, plain, ~ (op2(e21, e20) = op2(e22, e20)), inference(cnf_transformation, [], [f8])).
fof(f9949, plain, (~ spl80_113 | ~ spl80_38), inference(avatar_split_clause, [], [f9948, f1429, f1744])).
fof(f1429, plain, (spl80_38 <=> (e12 = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl80_38])])).
fof(f9948, plain, (~ (e12 = op1(e10, e12)) | ~ spl80_38), inference(forward_demodulation, [], [f394, f1431])).
fof(f1431, plain, ((e12 = op1(e13, e12)) | ~ spl80_38), inference(avatar_component_clause, [], [f1429])).
fof(f394, plain, ~ (op1(e10, e12) = op1(e13, e12)), inference(cnf_transformation, [], [f7])).
fof(f9928, plain, (~ spl80_89 | ~ spl80_90), inference(avatar_contradiction_clause, [], [f9927])).
fof(f9927, plain, ($false | (~ spl80_89 | ~ spl80_90)), inference(subsumption_resolution, [], [f9926, f581])).
fof(f581, plain, ~ (e13 = e14), inference(cnf_transformation, [], [f9])).
fof(f9926, plain, ((e13 = e14) | (~ spl80_89 | ~ spl80_90)), inference(forward_demodulation, [], [f1649, f1645])).
fof(f1649, plain, ((e14 = op1(e11, e12)) | ~ spl80_90), inference(avatar_component_clause, [], [f1647])).
fof(f1647, plain, (spl80_90 <=> (e14 = op1(e11, e12))), introduced(avatar_definition, [new_symbols(naming, [spl80_90])])).
fof(f9923, plain, (~ spl80_76 | ~ spl80_91), inference(avatar_split_clause, [], [f9830, f1652, f1589])).
fof(f1589, plain, (spl80_76 <=> (e10 = op1(e11, e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_76])])).
fof(f9830, plain, (~ (e10 = op1(e11, e14)) | ~ spl80_91), inference(backward_demodulation, [], [f438, f1654])).
fof(f438, plain, ~ (op1(e11, e11) = op1(e11, e14)), inference(cnf_transformation, [], [f7])).
fof(f9896, plain, (~ spl80_17 | ~ spl80_12), inference(avatar_split_clause, [], [f9492, f1320, f1341])).
fof(f1341, plain, (spl80_17 <=> (e11 = op1(e14, e11))), introduced(avatar_definition, [new_symbols(naming, [spl80_17])])).
fof(f9492, plain, (~ (e11 = op1(e14, e11)) | ~ spl80_12), inference(forward_demodulation, [], [f466, f1322])).
fof(f466, plain, ~ (op1(e14, e11) = op1(e14, e12)), inference(cnf_transformation, [], [f7])).
fof(f9875, plain, (~ spl80_53 | ~ spl80_54), inference(avatar_contradiction_clause, [], [f9874])).
fof(f9874, plain, ($false | (~ spl80_53 | ~ spl80_54)), inference(subsumption_resolution, [], [f9873, f579])).
fof(f9873, plain, ((e12 = e13) | (~ spl80_53 | ~ spl80_54)), inference(backward_demodulation, [], [f1498, f1494])).
fof(f1494, plain, ((e12 = op1(e12, e14)) | ~ spl80_53), inference(avatar_component_clause, [], [f1492])).
fof(f1492, plain, (spl80_53 <=> (e12 = op1(e12, e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_53])])).
fof(f9872, plain, (~ spl80_56 | ~ spl80_57), inference(avatar_contradiction_clause, [], [f9871])).
fof(f9871, plain, ($false | (~ spl80_56 | ~ spl80_57)), inference(subsumption_resolution, [], [f9870, f572])).
fof(f572, plain, ~ (e10 = e11), inference(cnf_transformation, [], [f9])).
fof(f9870, plain, ((e10 = e11) | (~ spl80_56 | ~ spl80_57)), inference(backward_demodulation, [], [f1511, f1507])).
fof(f1507, plain, ((e10 = op1(e12, e13)) | ~ spl80_56), inference(avatar_component_clause, [], [f1505])).
fof(f1505, plain, (spl80_56 <=> (e10 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl80_56])])).
fof(f9869, plain, (~ spl80_57 | ~ spl80_58), inference(avatar_contradiction_clause, [], [f9868])).
fof(f9868, plain, ($false | (~ spl80_57 | ~ spl80_58)), inference(subsumption_resolution, [], [f9867, f576])).
fof(f9867, plain, ((e11 = e12) | (~ spl80_57 | ~ spl80_58)), inference(backward_demodulation, [], [f1515, f1511])).
fof(f1515, plain, ((e12 = op1(e12, e13)) | ~ spl80_58), inference(avatar_component_clause, [], [f1513])).
fof(f1513, plain, (spl80_58 <=> (e12 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl80_58])])).
fof(f9853, plain, (~ spl80_68 | ~ spl80_73), inference(avatar_split_clause, [], [f9851, f1576, f1555])).
fof(f1555, plain, (spl80_68 <=> (e12 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl80_68])])).
fof(f9851, plain, (~ (e12 = op1(e12, e11)) | ~ spl80_73), inference(backward_demodulation, [], [f442, f1578])).
fof(f442, plain, ~ (op1(e12, e10) = op1(e12, e11)), inference(cnf_transformation, [], [f7])).
fof(f9836, plain, (~ spl80_16 | ~ spl80_91), inference(avatar_split_clause, [], [f9829, f1652, f1337])).
fof(f1337, plain, (spl80_16 <=> (e10 = op1(e14, e11))), introduced(avatar_definition, [new_symbols(naming, [spl80_16])])).
fof(f9829, plain, (~ (e10 = op1(e14, e11)) | ~ spl80_91), inference(backward_demodulation, [], [f388, f1654])).
fof(f388, plain, ~ (op1(e11, e11) = op1(e14, e11)), inference(cnf_transformation, [], [f7])).
fof(f9835, plain, (~ spl80_66 | ~ spl80_91), inference(avatar_split_clause, [], [f9827, f1652, f1547])).
fof(f1547, plain, (spl80_66 <=> (e10 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl80_66])])).
fof(f9827, plain, (~ (e10 = op1(e12, e11)) | ~ spl80_91), inference(backward_demodulation, [], [f386, f1654])).
fof(f386, plain, ~ (op1(e11, e11) = op1(e12, e11)), inference(cnf_transformation, [], [f7])).
fof(f9817, plain, (~ spl80_105 | ~ spl80_110), inference(avatar_split_clause, [], [f9816, f1731, f1710])).
fof(f1731, plain, (spl80_110 <=> (e14 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl80_110])])).
fof(f9816, plain, (~ (e14 = op1(e10, e14)) | ~ spl80_110), inference(backward_demodulation, [], [f431, f1733])).
fof(f1733, plain, ((e14 = op1(e10, e13)) | ~ spl80_110), inference(avatar_component_clause, [], [f1731])).
fof(f431, plain, ~ (op1(e10, e13) = op1(e10, e14)), inference(cnf_transformation, [], [f7])).
fof(f9814, plain, (~ spl80_103 | ~ spl80_113), inference(avatar_split_clause, [], [f9810, f1744, f1702])).
fof(f1702, plain, (spl80_103 <=> (e12 = op1(e10, e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_103])])).
fof(f9810, plain, (~ (e12 = op1(e10, e14)) | ~ spl80_113), inference(backward_demodulation, [], [f430, f1746])).
fof(f430, plain, ~ (op1(e10, e12) = op1(e10, e14)), inference(cnf_transformation, [], [f7])).
fof(f9813, plain, (~ spl80_108 | ~ spl80_113), inference(avatar_split_clause, [], [f9809, f1744, f1723])).
fof(f1723, plain, (spl80_108 <=> (e12 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl80_108])])).
fof(f9809, plain, (~ (e12 = op1(e10, e13)) | ~ spl80_113), inference(backward_demodulation, [], [f429, f1746])).
fof(f429, plain, ~ (op1(e10, e12) = op1(e10, e13)), inference(cnf_transformation, [], [f7])).
fof(f9794, plain, (~ spl80_12 | ~ spl80_128), inference(avatar_contradiction_clause, [], [f9793])).
fof(f9793, plain, ($false | (~ spl80_12 | ~ spl80_128)), inference(subsumption_resolution, [], [f9792, f578])).
fof(f9792, plain, ((e11 = e14) | (~ spl80_12 | ~ spl80_128)), inference(forward_demodulation, [], [f9775, f1322])).
fof(f9775, plain, ((e14 = op1(e14, e12)) | ~ spl80_128), inference(backward_demodulation, [], [f234, f1809])).
fof(f1809, plain, ((e12 = unit1) | ~ spl80_128), inference(avatar_component_clause, [], [f1807])).
fof(f1807, plain, (spl80_128 <=> (e12 = unit1)), introduced(avatar_definition, [new_symbols(naming, [spl80_128])])).
fof(f9747, plain, (~ spl80_156 | ~ spl80_157), inference(avatar_contradiction_clause, [], [f9746])).
fof(f9746, plain, ($false | (~ spl80_156 | ~ spl80_157)), inference(subsumption_resolution, [], [f9745, f582])).
fof(f9745, plain, ((e20 = e21) | (~ spl80_156 | ~ spl80_157)), inference(forward_demodulation, [], [f1981, f1977])).
fof(f1977, plain, ((e20 = op2(e23, e24)) | ~ spl80_156), inference(avatar_component_clause, [], [f1975])).
fof(f1975, plain, (spl80_156 <=> (e20 = op2(e23, e24))), introduced(avatar_definition, [new_symbols(naming, [spl80_156])])).
fof(f9731, plain, (~ spl80_183 | ~ spl80_184), inference(avatar_contradiction_clause, [], [f9730])).
fof(f9730, plain, ($false | (~ spl80_183 | ~ spl80_184)), inference(subsumption_resolution, [], [f9729, f589])).
fof(f9729, plain, ((e22 = e23) | (~ spl80_183 | ~ spl80_184)), inference(backward_demodulation, [], [f2094, f2090])).
fof(f2090, plain, ((e22 = op2(e22, e24)) | ~ spl80_183), inference(avatar_component_clause, [], [f2088])).
fof(f2088, plain, (spl80_183 <=> (e22 = op2(e22, e24))), introduced(avatar_definition, [new_symbols(naming, [spl80_183])])).
fof(f9709, plain, (~ spl80_203 | ~ spl80_205), inference(avatar_contradiction_clause, [], [f9708])).
fof(f9708, plain, ($false | (~ spl80_203 | ~ spl80_205)), inference(subsumption_resolution, [], [f9707, f590])).
fof(f590, plain, ~ (e22 = e24), inference(cnf_transformation, [], [f10])).
fof(f9707, plain, ((e22 = e24) | (~ spl80_203 | ~ spl80_205)), inference(backward_demodulation, [], [f2182, f2174])).
fof(f2182, plain, ((e24 = op2(e22, e20)) | ~ spl80_205), inference(avatar_component_clause, [], [f2180])).
fof(f2180, plain, (spl80_205 <=> (e24 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl80_205])])).
fof(f9691, plain, (spl80_424 | ~ spl80_228), inference(avatar_split_clause, [], [f9690, f2277, f3631])).
fof(f2277, plain, (spl80_228 <=> (e22 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl80_228])])).
fof(f9690, plain, ((e22 = h1(e11)) | ~ spl80_228), inference(backward_demodulation, [], [f676, f2279])).
fof(f2279, plain, ((e22 = op2(e21, e20)) | ~ spl80_228), inference(avatar_component_clause, [], [f2277])).
fof(f9684, plain, (spl80_792 | ~ spl80_245), inference(avatar_split_clause, [], [f9682, f2348, f6779])).
fof(f6779, plain, (spl80_792 <=> (e24 = h2(e13))), introduced(avatar_definition, [new_symbols(naming, [spl80_792])])).
fof(f2348, plain, (spl80_245 <=> (e24 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl80_245])])).
fof(f9682, plain, ((e24 = h2(e13)) | ~ spl80_245), inference(backward_demodulation, [], [f682, f2350])).
fof(f2350, plain, ((e24 = op2(e20, e22)) | ~ spl80_245), inference(avatar_component_clause, [], [f2348])).
fof(f9681, plain, (spl80_425 | ~ spl80_248), inference(avatar_split_clause, [], [f9680, f2361, f3638])).
fof(f2361, plain, (spl80_248 <=> (e22 = op2(e20, e21))), introduced(avatar_definition, [new_symbols(naming, [spl80_248])])).
fof(f9680, plain, ((e22 = h1(e13)) | ~ spl80_248), inference(backward_demodulation, [], [f677, f2363])).
fof(f2363, plain, ((e22 = op2(e20, e21)) | ~ spl80_248), inference(avatar_component_clause, [], [f2361])).
fof(f9679, plain, (~ spl80_142 | ~ spl80_258), inference(avatar_contradiction_clause, [], [f9678])).
fof(f9678, plain, ($false | (~ spl80_142 | ~ spl80_258)), inference(subsumption_resolution, [], [f9677, f588])).
fof(f9677, plain, ((e21 = e24) | (~ spl80_142 | ~ spl80_258)), inference(forward_demodulation, [], [f9666, f1918])).
fof(f9666, plain, ((e24 = op2(e24, e22)) | ~ spl80_258), inference(backward_demodulation, [], [f320, f2405])).
fof(f2405, plain, ((e22 = unit2) | ~ spl80_258), inference(avatar_component_clause, [], [f2403])).
fof(f2403, plain, (spl80_258 <=> (e22 = unit2)), introduced(avatar_definition, [new_symbols(naming, [spl80_258])])).
fof(f9602, plain, (spl80_295 | ~ spl80_206), inference(avatar_split_clause, [], [f9601, f2185, f2852])).
fof(f2185, plain, (spl80_206 <=> (e20 = op2(e21, e24))), introduced(avatar_definition, [new_symbols(naming, [spl80_206])])).
fof(f9601, plain, ((e20 = h8(e13)) | ~ spl80_206), inference(forward_demodulation, [], [f712, f2187])).
fof(f2187, plain, ((e20 = op2(e21, e24)) | ~ spl80_206), inference(avatar_component_clause, [], [f2185])).
fof(f9581, plain, (~ spl80_371 | ~ spl80_169), inference(avatar_split_clause, [], [f9580, f2029, f3314])).
fof(f2029, plain, (spl80_169 <=> (e23 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl80_169])])).
fof(f9580, plain, (~ (e23 = h6(e13)) | ~ spl80_169), inference(forward_demodulation, [], [f2569, f2031])).
fof(f2031, plain, ((e23 = op2(e23, e22)) | ~ spl80_169), inference(avatar_component_clause, [], [f2029])).
fof(f2569, plain, ~ (op2(e23, e22) = h6(e13)), inference(backward_demodulation, [], [f497, f702])).
fof(f497, plain, ~ (op2(e21, e22) = op2(e23, e22)), inference(cnf_transformation, [], [f8])).
fof(f9539, plain, (spl80_270 | ~ spl80_136), inference(avatar_split_clause, [], [f9538, f1891, f2705])).
fof(f2705, plain, (spl80_270 <=> (e20 = h16(e11))), introduced(avatar_definition, [new_symbols(naming, [spl80_270])])).
fof(f1891, plain, (spl80_136 <=> (e20 = op2(e24, e23))), introduced(avatar_definition, [new_symbols(naming, [spl80_136])])).
fof(f9538, plain, ((e20 = h16(e11)) | ~ spl80_136), inference(forward_demodulation, [], [f751, f1893])).
fof(f1893, plain, ((e20 = op2(e24, e23)) | ~ spl80_136), inference(avatar_component_clause, [], [f1891])).
fof(f751, plain, (op2(e24, e23) = h16(e11)), inference(cnf_transformation, [], [f31])).
fof(f31, plain, ((op2(e23, e24) = h16(e13)) & (op2(e24, e23) = h16(e11)) & (h16(e10) = op2(op2(e24, e23), op2(e24, e23))) & (e24 = h16(e14)) & (e23 = h16(e12))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG049+1.p', ax31)).
fof(f9524, plain, (spl80_91 | ~ spl80_12), inference(avatar_split_clause, [], [f9006, f1320, f1652])).
fof(f9006, plain, ((e10 = op1(e11, e11)) | ~ spl80_12), inference(forward_demodulation, [], [f667, f1322])).
fof(f667, plain, (e10 = op1(op1(e14, e12), op1(e14, e12))), inference(cnf_transformation, [], [f14])).
fof(f14, plain, ((e13 = op1(e12, e14)) & (e11 = op1(e14, e12)) & (e10 = op1(op1(e14, e12), op1(e14, e12)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG049+1.p', ax14)).
fof(f9519, plain, (~ spl80_79 | ~ spl80_54), inference(avatar_split_clause, [], [f9075, f1496, f1601])).
fof(f1601, plain, (spl80_79 <=> (e13 = op1(e11, e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_79])])).
fof(f9075, plain, (~ (e13 = op1(e11, e14)) | ~ spl80_54), inference(forward_demodulation, [], [f416, f1498])).
fof(f416, plain, ~ (op1(e11, e14) = op1(e12, e14)), inference(cnf_transformation, [], [f7])).
fof(f9517, plain, (~ spl80_77 | ~ spl80_97), inference(avatar_split_clause, [], [f9516, f1677, f1593])).
fof(f1593, plain, (spl80_77 <=> (e11 = op1(e11, e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_77])])).
fof(f9516, plain, (~ (e11 = op1(e11, e14)) | ~ spl80_97), inference(forward_demodulation, [], [f435, f1679])).
fof(f435, plain, ~ (op1(e11, e10) = op1(e11, e14)), inference(cnf_transformation, [], [f7])).
fof(f9511, plain, (~ spl80_70 | ~ spl80_60), inference(avatar_split_clause, [], [f9510, f1521, f1563])).
fof(f1521, plain, (spl80_60 <=> (e14 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl80_60])])).
fof(f9510, plain, (~ (e14 = op1(e12, e11)) | ~ spl80_60), inference(forward_demodulation, [], [f447, f1523])).
fof(f1523, plain, ((e14 = op1(e12, e13)) | ~ spl80_60), inference(avatar_component_clause, [], [f1521])).
fof(f447, plain, ~ (op1(e12, e11) = op1(e12, e13)), inference(cnf_transformation, [], [f7])).
fof(f9478, plain, (~ spl80_2 | ~ spl80_12), inference(avatar_split_clause, [], [f9477, f1320, f1278])).
fof(f1278, plain, (spl80_2 <=> (e11 = op1(e14, e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_2])])).
fof(f9477, plain, (~ (e11 = op1(e14, e14)) | ~ spl80_12), inference(forward_demodulation, [], [f470, f1322])).
fof(f470, plain, ~ (op1(e14, e12) = op1(e14, e14)), inference(cnf_transformation, [], [f7])).
fof(f9428, plain, (~ spl80_82 | ~ spl80_97), inference(avatar_split_clause, [], [f9425, f1677, f1614])).
fof(f9425, plain, (~ (e11 = op1(e11, e13)) | ~ spl80_97), inference(backward_demodulation, [], [f434, f1679])).
fof(f434, plain, ~ (op1(e11, e10) = op1(e11, e13)), inference(cnf_transformation, [], [f7])).
fof(f9426, plain, (~ spl80_72 | ~ spl80_97), inference(avatar_split_clause, [], [f9423, f1677, f1572])).
fof(f1572, plain, (spl80_72 <=> (e11 = op1(e12, e10))), introduced(avatar_definition, [new_symbols(naming, [spl80_72])])).
fof(f9423, plain, (~ (e11 = op1(e12, e10)) | ~ spl80_97), inference(backward_demodulation, [], [f376, f1679])).
fof(f376, plain, ~ (op1(e11, e10) = op1(e12, e10)), inference(cnf_transformation, [], [f7])).
fof(f9420, plain, (~ spl80_84 | ~ spl80_109), inference(avatar_split_clause, [], [f9418, f1727, f1622])).
fof(f1622, plain, (spl80_84 <=> (e13 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl80_84])])).
fof(f9418, plain, (~ (e13 = op1(e11, e13)) | ~ spl80_109), inference(backward_demodulation, [], [f402, f1729])).
fof(f402, plain, ~ (op1(e10, e13) = op1(e11, e13)), inference(cnf_transformation, [], [f7])).
fof(f9389, plain, (spl80_96 | ~ spl80_127), inference(avatar_split_clause, [], [f9379, f1803, f1673])).
fof(f1673, plain, (spl80_96 <=> (e10 = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl80_96])])).
fof(f1803, plain, (spl80_127 <=> (e11 = unit1)), introduced(avatar_definition, [new_symbols(naming, [spl80_127])])).
fof(f9379, plain, ((e10 = op1(e11, e10)) | ~ spl80_127), inference(backward_demodulation, [], [f225, f1805])).
fof(f1805, plain, ((e11 = unit1) | ~ spl80_127), inference(avatar_component_clause, [], [f1803])).
fof(f225, plain, (e10 = op1(unit1, e10)), inference(cnf_transformation, [], [f2])).
fof(f9301, plain, (~ spl80_181 | ~ spl80_184), inference(avatar_contradiction_clause, [], [f9300])).
fof(f9300, plain, ($false | (~ spl80_181 | ~ spl80_184)), inference(subsumption_resolution, [], [f9299, f584])).
fof(f584, plain, ~ (e20 = e23), inference(cnf_transformation, [], [f10])).
fof(f9299, plain, ((e20 = e23) | (~ spl80_181 | ~ spl80_184)), inference(backward_demodulation, [], [f2094, f2082])).
fof(f2082, plain, ((e20 = op2(e22, e24)) | ~ spl80_181), inference(avatar_component_clause, [], [f2080])).
fof(f2080, plain, (spl80_181 <=> (e20 = op2(e22, e24))), introduced(avatar_definition, [new_symbols(naming, [spl80_181])])).
fof(f9272, plain, (~ spl80_213 | ~ spl80_215), inference(avatar_contradiction_clause, [], [f9271])).
fof(f9271, plain, ($false | (~ spl80_213 | ~ spl80_215)), inference(subsumption_resolution, [], [f9270, f590])).
fof(f9270, plain, ((e22 = e24) | (~ spl80_213 | ~ spl80_215)), inference(backward_demodulation, [], [f2224, f2216])).
fof(f2216, plain, ((e22 = op2(e21, e23)) | ~ spl80_213), inference(avatar_component_clause, [], [f2214])).
fof(f2214, plain, (spl80_213 <=> (e22 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl80_213])])).
fof(f9260, plain, (~ spl80_217 | spl80_414), inference(avatar_contradiction_clause, [], [f9259])).
fof(f9259, plain, ($false | (~ spl80_217 | spl80_414)), inference(subsumption_resolution, [], [f9256, f3574])).
fof(f3574, plain, (~ (e21 = h6(e13)) | spl80_414), inference(avatar_component_clause, [], [f3572])).
fof(f3572, plain, (spl80_414 <=> (e21 = h6(e13))), introduced(avatar_definition, [new_symbols(naming, [spl80_414])])).
fof(f9256, plain, ((e21 = h6(e13)) | ~ spl80_217), inference(backward_demodulation, [], [f702, f2233])).
fof(f2233, plain, ((e21 = op2(e21, e22)) | ~ spl80_217), inference(avatar_component_clause, [], [f2231])).
fof(f2231, plain, (spl80_217 <=> (e21 = op2(e21, e22))), introduced(avatar_definition, [new_symbols(naming, [spl80_217])])).
fof(f9255, plain, (~ spl80_227 | spl80_428), inference(avatar_contradiction_clause, [], [f9254])).
fof(f9254, plain, ($false | (~ spl80_227 | spl80_428)), inference(subsumption_resolution, [], [f9253, f3656])).
fof(f3656, plain, (~ (e21 = h1(e11)) | spl80_428), inference(avatar_component_clause, [], [f3654])).
fof(f9253, plain, ((e21 = h1(e11)) | ~ spl80_227), inference(backward_demodulation, [], [f676, f2275])).
fof(f9238, plain, (~ spl80_247 | spl80_429), inference(avatar_contradiction_clause, [], [f9237])).
fof(f9237, plain, ($false | (~ spl80_247 | spl80_429)), inference(subsumption_resolution, [], [f9236, f3664])).
fof(f3664, plain, (~ (e21 = h1(e13)) | spl80_429), inference(avatar_component_clause, [], [f3662])).
fof(f9236, plain, ((e21 = h1(e13)) | ~ spl80_247), inference(backward_demodulation, [], [f677, f2359])).
fof(f9199, plain, (spl80_287 | ~ spl80_209), inference(avatar_split_clause, [], [f9198, f2197, f2805])).
fof(f2805, plain, (spl80_287 <=> (e23 = h8(e13))), introduced(avatar_definition, [new_symbols(naming, [spl80_287])])).
fof(f2197, plain, (spl80_209 <=> (e23 = op2(e21, e24))), introduced(avatar_definition, [new_symbols(naming, [spl80_209])])).
fof(f9198, plain, ((e23 = h8(e13)) | ~ spl80_209), inference(forward_demodulation, [], [f712, f2199])).
fof(f2199, plain, ((e23 = op2(e21, e24)) | ~ spl80_209), inference(avatar_component_clause, [], [f2197])).
fof(f9183, plain, (~ spl80_370 | ~ spl80_184), inference(avatar_split_clause, [], [f9182, f2092, f3307])).
fof(f9182, plain, (~ (e23 = h6(e11)) | ~ spl80_184), inference(forward_demodulation, [], [f2579, f2094])).
fof(f2579, plain, ~ (op2(e22, e24) = h6(e11)), inference(backward_demodulation, [], [f548, f701])).
fof(f548, plain, ~ (op2(e22, e21) = op2(e22, e24)), inference(cnf_transformation, [], [f8])).
fof(f9181, plain, (~ spl80_287 | ~ spl80_184), inference(avatar_split_clause, [], [f9180, f2092, f2805])).
fof(f9180, plain, (~ (e23 = h8(e13)) | ~ spl80_184), inference(forward_demodulation, [], [f2592, f2094])).
fof(f2592, plain, ~ (op2(e22, e24) = h8(e13)), inference(backward_demodulation, [], [f516, f712])).
fof(f516, plain, ~ (op2(e21, e24) = op2(e22, e24)), inference(cnf_transformation, [], [f8])).
fof(f9152, plain, (spl80_286 | ~ spl80_149), inference(avatar_split_clause, [], [f9151, f1945, f2798])).
fof(f1945, plain, (spl80_149 <=> (e23 = op2(e24, e21))), introduced(avatar_definition, [new_symbols(naming, [spl80_149])])).
fof(f9151, plain, ((e23 = h8(e11)) | ~ spl80_149), inference(forward_demodulation, [], [f711, f1947])).
fof(f1947, plain, ((e23 = op2(e24, e21)) | ~ spl80_149), inference(avatar_component_clause, [], [f1945])).
fof(f9142, plain, (spl80_278 | ~ spl80_142), inference(avatar_split_clause, [], [f9141, f1916, f2752])).
fof(f9141, plain, ((e21 = h12(e11)) | ~ spl80_142), inference(forward_demodulation, [], [f731, f1918])).
fof(f9136, plain, (~ spl80_414 | ~ spl80_142), inference(avatar_split_clause, [], [f9135, f1916, f3572])).
fof(f9135, plain, (~ (e21 = h6(e13)) | ~ spl80_142), inference(forward_demodulation, [], [f2570, f1918])).
fof(f2570, plain, ~ (op2(e24, e22) = h6(e13)), inference(backward_demodulation, [], [f498, f702])).
fof(f498, plain, ~ (op2(e21, e22) = op2(e24, e22)), inference(cnf_transformation, [], [f8])).
fof(f9118, plain, (~ spl80_278 | ~ spl80_137), inference(avatar_split_clause, [], [f9117, f1895, f2752])).
fof(f1895, plain, (spl80_137 <=> (e21 = op2(e24, e23))), introduced(avatar_definition, [new_symbols(naming, [spl80_137])])).
fof(f9117, plain, (~ (e21 = h12(e11)) | ~ spl80_137), inference(forward_demodulation, [], [f2620, f1897])).
fof(f1897, plain, ((e21 = op2(e24, e23)) | ~ spl80_137), inference(avatar_component_clause, [], [f1895])).
fof(f2620, plain, ~ (op2(e24, e23) = h12(e11)), inference(backward_demodulation, [], [f569, f731])).
fof(f569, plain, ~ (op2(e24, e22) = op2(e24, e23)), inference(cnf_transformation, [], [f8])).
fof(f9046, plain, (~ spl80_43 | ~ spl80_18), inference(avatar_split_clause, [], [f9045, f1345, f1450])).
fof(f1345, plain, (spl80_18 <=> (e12 = op1(e14, e11))), introduced(avatar_definition, [new_symbols(naming, [spl80_18])])).
fof(f9045, plain, (~ (e12 = op1(e13, e11)) | ~ spl80_18), inference(forward_demodulation, [], [f391, f1347])).
fof(f1347, plain, ((e12 = op1(e14, e11)) | ~ spl80_18), inference(avatar_component_clause, [], [f1345])).
fof(f391, plain, ~ (op1(e13, e11) = op1(e14, e11)), inference(cnf_transformation, [], [f7])).
fof(f8955, plain, (~ spl80_10 | ~ spl80_25), inference(avatar_split_clause, [], [f8951, f1374, f1311])).
fof(f1311, plain, (spl80_10 <=> (e14 = op1(e14, e13))), introduced(avatar_definition, [new_symbols(naming, [spl80_10])])).
fof(f8951, plain, (~ (e14 = op1(e14, e13)) | ~ spl80_25), inference(backward_demodulation, [], [f464, f1376])).
fof(f464, plain, ~ (op1(e14, e10) = op1(e14, e13)), inference(cnf_transformation, [], [f7])).
fof(f8954, plain, (~ spl80_15 | ~ spl80_25), inference(avatar_split_clause, [], [f8950, f1374, f1332])).
fof(f1332, plain, (spl80_15 <=> (e14 = op1(e14, e12))), introduced(avatar_definition, [new_symbols(naming, [spl80_15])])).
fof(f8950, plain, (~ (e14 = op1(e14, e12)) | ~ spl80_25), inference(backward_demodulation, [], [f463, f1376])).
fof(f463, plain, ~ (op1(e14, e10) = op1(e14, e12)), inference(cnf_transformation, [], [f7])).
fof(f8953, plain, (~ spl80_20 | ~ spl80_25), inference(avatar_split_clause, [], [f8949, f1374, f1353])).
fof(f1353, plain, (spl80_20 <=> (e14 = op1(e14, e11))), introduced(avatar_definition, [new_symbols(naming, [spl80_20])])).
fof(f8949, plain, (~ (e14 = op1(e14, e11)) | ~ spl80_25), inference(backward_demodulation, [], [f462, f1376])).
fof(f462, plain, ~ (op1(e14, e10) = op1(e14, e11)), inference(cnf_transformation, [], [f7])).
fof(f8948, plain, (~ spl80_26 | ~ spl80_27), inference(avatar_contradiction_clause, [], [f8947])).
fof(f8947, plain, ($false | (~ spl80_26 | ~ spl80_27)), inference(subsumption_resolution, [], [f8946, f572])).
fof(f8946, plain, ((e10 = e11) | (~ spl80_26 | ~ spl80_27)), inference(backward_demodulation, [], [f1385, f1381])).
fof(f1381, plain, ((e10 = op1(e13, e14)) | ~ spl80_26), inference(avatar_component_clause, [], [f1379])).
fof(f1379, plain, (spl80_26 <=> (e10 = op1(e13, e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_26])])).
fof(f8945, plain, (~ spl80_27 | ~ spl80_28), inference(avatar_contradiction_clause, [], [f8944])).
fof(f8944, plain, ($false | (~ spl80_27 | ~ spl80_28)), inference(subsumption_resolution, [], [f8943, f576])).
fof(f8943, plain, ((e11 = e12) | (~ spl80_27 | ~ spl80_28)), inference(backward_demodulation, [], [f1389, f1385])).
fof(f1389, plain, ((e12 = op1(e13, e14)) | ~ spl80_28), inference(avatar_component_clause, [], [f1387])).
fof(f1387, plain, (spl80_28 <=> (e12 = op1(e13, e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_28])])).
fof(f8907, plain, (~ spl80_52 | ~ spl80_54), inference(avatar_contradiction_clause, [], [f8906])).
fof(f8906, plain, ($false | (~ spl80_52 | ~ spl80_54)), inference(subsumption_resolution, [], [f8905, f577])).
fof(f8905, plain, ((e11 = e13) | (~ spl80_52 | ~ spl80_54)), inference(backward_demodulation, [], [f1498, f1490])).
fof(f1490, plain, ((e11 = op1(e12, e14)) | ~ spl80_52), inference(avatar_component_clause, [], [f1488])).
fof(f1488, plain, (spl80_52 <=> (e11 = op1(e12, e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_52])])).
fof(f8809, plain, (~ spl80_91 | ~ spl80_96), inference(avatar_split_clause, [], [f8802, f1673, f1652])).
fof(f8802, plain, (~ (e10 = op1(e11, e11)) | ~ spl80_96), inference(backward_demodulation, [], [f432, f1675])).
fof(f1675, plain, ((e10 = op1(e11, e10)) | ~ spl80_96), inference(avatar_component_clause, [], [f1673])).
fof(f432, plain, ~ (op1(e11, e10) = op1(e11, e11)), inference(cnf_transformation, [], [f7])).
fof(f8743, plain, (~ spl80_101 | ~ spl80_121), inference(avatar_split_clause, [], [f8717, f1778, f1694])).
fof(f1694, plain, (spl80_101 <=> (e10 = op1(e10, e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_101])])).
fof(f8717, plain, (~ (e10 = op1(e10, e14)) | ~ spl80_121), inference(backward_demodulation, [], [f425, f1780])).
fof(f425, plain, ~ (op1(e10, e10) = op1(e10, e14)), inference(cnf_transformation, [], [f7])).
fof(f8738, plain, (~ spl80_46 | ~ spl80_121), inference(avatar_split_clause, [], [f8712, f1778, f1463])).
fof(f1463, plain, (spl80_46 <=> (e10 = op1(e13, e10))), introduced(avatar_definition, [new_symbols(naming, [spl80_46])])).
fof(f8712, plain, (~ (e10 = op1(e13, e10)) | ~ spl80_121), inference(backward_demodulation, [], [f374, f1780])).
fof(f374, plain, ~ (op1(e10, e10) = op1(e13, e10)), inference(cnf_transformation, [], [f7])).
fof(f8737, plain, (~ spl80_71 | ~ spl80_121), inference(avatar_split_clause, [], [f8711, f1778, f1568])).
fof(f1568, plain, (spl80_71 <=> (e10 = op1(e12, e10))), introduced(avatar_definition, [new_symbols(naming, [spl80_71])])).
fof(f8711, plain, (~ (e10 = op1(e12, e10)) | ~ spl80_121), inference(backward_demodulation, [], [f373, f1780])).
fof(f373, plain, ~ (op1(e10, e10) = op1(e12, e10)), inference(cnf_transformation, [], [f7])).
fof(f8708, plain, (spl80_105 | ~ spl80_126), inference(avatar_split_clause, [], [f8698, f1799, f1710])).
fof(f8698, plain, ((e14 = op1(e10, e14)) | ~ spl80_126), inference(backward_demodulation, [], [f233, f1801])).
fof(f8707, plain, (spl80_49 | ~ spl80_126), inference(avatar_split_clause, [], [f8697, f1799, f1475])).
fof(f8697, plain, ((e13 = op1(e13, e10)) | ~ spl80_126), inference(backward_demodulation, [], [f232, f1801])).
fof(f8701, plain, (spl80_121 | ~ spl80_126), inference(avatar_split_clause, [], [f8691, f1799, f1778])).
fof(f8691, plain, ((e10 = op1(e10, e10)) | ~ spl80_126), inference(backward_demodulation, [], [f226, f1801])).
fof(f226, plain, (e10 = op1(e10, unit1)), inference(cnf_transformation, [], [f2])).
fof(f8689, plain, (~ spl80_270 | ~ spl80_131), inference(avatar_split_clause, [], [f8681, f1870, f2705])).
fof(f8681, plain, (~ (e20 = h16(e11)) | ~ spl80_131), inference(backward_demodulation, [], [f2637, f1872])).
fof(f2637, plain, ~ (op2(e24, e24) = h16(e11)), inference(backward_demodulation, [], [f571, f751])).
fof(f571, plain, ~ (op2(e24, e23) = op2(e24, e24)), inference(cnf_transformation, [], [f8])).
fof(f8635, plain, (~ spl80_148 | ~ spl80_149), inference(avatar_contradiction_clause, [], [f8634])).
fof(f8634, plain, ($false | (~ spl80_148 | ~ spl80_149)), inference(subsumption_resolution, [], [f8633, f589])).
fof(f8633, plain, ((e22 = e23) | (~ spl80_148 | ~ spl80_149)), inference(backward_demodulation, [], [f1947, f1943])).
fof(f1943, plain, ((e22 = op2(e24, e21)) | ~ spl80_148), inference(avatar_component_clause, [], [f1941])).
fof(f1941, plain, (spl80_148 <=> (e22 = op2(e24, e21))), introduced(avatar_definition, [new_symbols(naming, [spl80_148])])).
fof(f8541, plain, (~ spl80_182 | ~ spl80_184), inference(avatar_contradiction_clause, [], [f8540])).
fof(f8540, plain, ($false | (~ spl80_182 | ~ spl80_184)), inference(subsumption_resolution, [], [f8539, f587])).
fof(f8539, plain, ((e21 = e23) | (~ spl80_182 | ~ spl80_184)), inference(backward_demodulation, [], [f2094, f2086])).
fof(f2086, plain, ((e21 = op2(e22, e24)) | ~ spl80_182), inference(avatar_component_clause, [], [f2084])).
fof(f2084, plain, (spl80_182 <=> (e21 = op2(e22, e24))), introduced(avatar_definition, [new_symbols(naming, [spl80_182])])).
fof(f8538, plain, (~ spl80_184 | ~ spl80_185), inference(avatar_contradiction_clause, [], [f8537])).
fof(f8537, plain, ($false | (~ spl80_184 | ~ spl80_185)), inference(subsumption_resolution, [], [f8536, f591])).
fof(f8536, plain, ((e23 = e24) | (~ spl80_184 | ~ spl80_185)), inference(backward_demodulation, [], [f2098, f2094])).
fof(f2098, plain, ((e24 = op2(e22, e24)) | ~ spl80_185), inference(avatar_component_clause, [], [f2096])).
fof(f2096, plain, (spl80_185 <=> (e24 = op2(e22, e24))), introduced(avatar_definition, [new_symbols(naming, [spl80_185])])).
fof(f8493, plain, (spl80_378 | ~ spl80_196), inference(avatar_split_clause, [], [f8490, f2143, f3355])).
fof(f2143, plain, (spl80_196 <=> (e20 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl80_196])])).
fof(f8490, plain, ((e20 = h6(e11)) | ~ spl80_196), inference(backward_demodulation, [], [f701, f2145])).
fof(f2145, plain, ((e20 = op2(e22, e21)) | ~ spl80_196), inference(avatar_component_clause, [], [f2143])).
fof(f8488, plain, (spl80_455 | ~ spl80_201), inference(avatar_split_clause, [], [f8486, f2164, f3822])).
fof(f3822, plain, (spl80_455 <=> (e20 = h2(e11))), introduced(avatar_definition, [new_symbols(naming, [spl80_455])])).
fof(f2164, plain, (spl80_201 <=> (e20 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl80_201])])).
fof(f8486, plain, ((e20 = h2(e11)) | ~ spl80_201), inference(backward_demodulation, [], [f681, f2166])).
fof(f2166, plain, ((e20 = op2(e22, e20)) | ~ spl80_201), inference(avatar_component_clause, [], [f2164])).
fof(f8477, plain, (~ spl80_575 | ~ spl80_210), inference(avatar_split_clause, [], [f8472, f2201, f4970])).
fof(f4970, plain, (spl80_575 <=> (e24 = h4(e13))), introduced(avatar_definition, [new_symbols(naming, [spl80_575])])).
fof(f2201, plain, (spl80_210 <=> (e24 = op2(e21, e24))), introduced(avatar_definition, [new_symbols(naming, [spl80_210])])).
fof(f8472, plain, (~ (e24 = h4(e13)) | ~ spl80_210), inference(backward_demodulation, [], [f2554, f2203])).
fof(f2203, plain, ((e24 = op2(e21, e24)) | ~ spl80_210), inference(avatar_component_clause, [], [f2201])).
fof(f2554, plain, ~ (op2(e21, e24) = h4(e13)), inference(backward_demodulation, [], [f512, f692])).
fof(f512, plain, ~ (op2(e20, e24) = op2(e21, e24)), inference(cnf_transformation, [], [f8])).
fof(f8466, plain, (spl80_341 | ~ spl80_211), inference(avatar_split_clause, [], [f8462, f2206, f3125])).
fof(f2206, plain, (spl80_211 <=> (e20 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl80_211])])).
fof(f8462, plain, ((e20 = h7(e13)) | ~ spl80_211), inference(backward_demodulation, [], [f707, f2208])).
fof(f2208, plain, ((e20 = op2(e21, e23)) | ~ spl80_211), inference(avatar_component_clause, [], [f2206])).
fof(f8424, plain, (~ spl80_792 | ~ spl80_235), inference(avatar_split_clause, [], [f8420, f2306, f6779])).
fof(f8420, plain, (~ (e24 = h2(e13)) | ~ spl80_235), inference(backward_demodulation, [], [f2532, f2308])).
fof(f2532, plain, ~ (op2(e20, e24) = h2(e13)), inference(backward_demodulation, [], [f530, f682])).
fof(f530, plain, ~ (op2(e20, e22) = op2(e20, e24)), inference(cnf_transformation, [], [f8])).
fof(f8422, plain, (spl80_575 | ~ spl80_235), inference(avatar_split_clause, [], [f8418, f2306, f4970])).
fof(f8418, plain, ((e24 = h4(e13)) | ~ spl80_235), inference(backward_demodulation, [], [f692, f2308])).
fof(f8415, plain, (spl80_446 | ~ spl80_236), inference(avatar_split_clause, [], [f8412, f2311, f3766])).
fof(f3766, plain, (spl80_446 <=> (e20 = h3(e13))), introduced(avatar_definition, [new_symbols(naming, [spl80_446])])).
fof(f2311, plain, (spl80_236 <=> (e20 = op2(e20, e23))), introduced(avatar_definition, [new_symbols(naming, [spl80_236])])).
fof(f8412, plain, ((e20 = h3(e13)) | ~ spl80_236), inference(backward_demodulation, [], [f687, f2313])).
fof(f2313, plain, ((e20 = op2(e20, e23)) | ~ spl80_236), inference(avatar_component_clause, [], [f2311])).
fof(f8404, plain, (~ spl80_438 | ~ spl80_251), inference(avatar_split_clause, [], [f8396, f2374, f3717])).
fof(f8396, plain, (~ (e20 = h4(e13)) | ~ spl80_251), inference(backward_demodulation, [], [f2558, f2376])).
fof(f2558, plain, ~ (op2(e20, e20) = h4(e13)), inference(backward_demodulation, [], [f525, f692])).
fof(f525, plain, ~ (op2(e20, e20) = op2(e20, e24)), inference(cnf_transformation, [], [f8])).
fof(f8402, plain, (~ spl80_446 | ~ spl80_251), inference(avatar_split_clause, [], [f8394, f2374, f3766])).
fof(f8394, plain, (~ (e20 = h3(e13)) | ~ spl80_251), inference(backward_demodulation, [], [f2545, f2376])).
fof(f2545, plain, ~ (op2(e20, e20) = h3(e13)), inference(backward_demodulation, [], [f524, f687])).
fof(f524, plain, ~ (op2(e20, e20) = op2(e20, e23)), inference(cnf_transformation, [], [f8])).
fof(f8401, plain, (~ spl80_455 | ~ spl80_251), inference(avatar_split_clause, [], [f8393, f2374, f3822])).
fof(f8393, plain, (~ (e20 = h2(e11)) | ~ spl80_251), inference(backward_demodulation, [], [f2533, f2376])).
fof(f2533, plain, ~ (op2(e20, e20) = h2(e11)), inference(backward_demodulation, [], [f473, f681])).
fof(f473, plain, ~ (op2(e20, e20) = op2(e22, e20)), inference(cnf_transformation, [], [f8])).
fof(f8385, plain, (spl80_239 | ~ spl80_256), inference(avatar_split_clause, [], [f8375, f2395, f2323])).
fof(f8375, plain, ((e23 = op2(e20, e23)) | ~ spl80_256), inference(backward_demodulation, [], [f317, f2397])).
fof(f317, plain, (e23 = op2(unit2, e23)), inference(cnf_transformation, [], [f5])).
fof(f8384, plain, (spl80_203 | ~ spl80_256), inference(avatar_split_clause, [], [f8374, f2395, f2172])).
fof(f8374, plain, ((e22 = op2(e22, e20)) | ~ spl80_256), inference(backward_demodulation, [], [f316, f2397])).
fof(f316, plain, (e22 = op2(e22, unit2)), inference(cnf_transformation, [], [f5])).
fof(f8383, plain, (spl80_243 | ~ spl80_256), inference(avatar_split_clause, [], [f8373, f2395, f2340])).
fof(f8373, plain, ((e22 = op2(e20, e22)) | ~ spl80_256), inference(backward_demodulation, [], [f315, f2397])).
fof(f4641, plain, (~ spl80_521 | spl80_293 | spl80_289 | spl80_285 | ~ spl80_523 | ~ spl80_524 | ~ spl80_525 | ~ spl80_526 | ~ spl80_527 | ~ spl80_528 | ~ spl80_529 | ~ spl80_530 | ~ spl80_531 | ~ spl80_532 | ~ spl80_533 | ~ spl80_534 | ~ spl80_535 | ~ spl80_536 | ~ spl80_537 | ~ spl80_538 | ~ spl80_539 | ~ spl80_540 | ~ spl80_541 | ~ spl80_542 | ~ spl80_543 | ~ spl80_544 | ~ spl80_545 | ~ spl80_546), inference(avatar_split_clause, [], [f4640, f4595, f4591, f4587, f4583, f4579, f4575, f4571, f4567, f4563, f4559, f4555, f4551, f4547, f4543, f4539, f4535, f4531, f4527, f4523, f4519, f4515, f4511, f4507, f4503, f2794, f2817, f2841, f4495])).
fof(f2841, plain, (spl80_293 <=> sP68), introduced(avatar_definition, [new_symbols(naming, [spl80_293])])).
fof(f2817, plain, (spl80_289 <=> sP70), introduced(avatar_definition, [new_symbols(naming, [spl80_289])])).
fof(f2794, plain, (spl80_285 <=> sP71), introduced(avatar_definition, [new_symbols(naming, [spl80_285])])).
fof(f4640, plain, (~ (h18(op1(e10, e11)) = op2(h18(e10), h8(e13))) | ~ (h18(op1(e10, e12)) = op2(h18(e10), e24)) | ~ (h18(op1(e10, e13)) = op2(h18(e10), h8(e11))) | ~ (h18(op1(e10, e14)) = op2(h18(e10), e21)) | ~ (h18(op1(e11, e10)) = op2(h8(e13), h18(e10))) | ~ (h18(e10) = h18(op1(e11, e11))) | ~ (h18(op1(e11, e12)) = op2(h8(e13), e24)) | ~ (op2(h8(e13), h8(e11)) = h18(op1(e11, e13))) | ~ (h18(op1(e11, e14)) = op2(h8(e13), e21)) | ~ (h18(op1(e12, e10)) = op2(e24, h18(e10))) | ~ (h18(op1(e12, e11)) = op2(e24, h8(e13))) | ~ (op2(e24, e24) = h18(op1(e12, e12))) | ~ (h18(op1(e12, e13)) = op2(e24, h8(e11))) | ~ (h8(e11) = h18(op1(e12, e14))) | ~ (h18(op1(e13, e10)) = op2(h8(e11), h18(e10))) | ~ (op2(h8(e11), h8(e13)) = h18(op1(e13, e11))) | ~ (h18(op1(e13, e12)) = op2(h8(e11), e24)) | ~ (h8(e10) = h18(op1(e13, e13))) | ~ (h18(op1(e13, e14)) = op2(h8(e11), e21)) | ~ (h18(op1(e14, e10)) = op2(e21, h18(e10))) | ~ (h18(op1(e14, e11)) = op2(e21, h8(e13))) | ~ (h8(e13) = h18(op1(e14, e12))) | ~ (h18(op1(e14, e13)) = op2(e21, h8(e11))) | ~ (op2(e21, e21) = h18(op1(e14, e14))) | sP71 | sP70 | sP68 | ~ (h18(op1(e10, e10)) = op2(h18(e10), h18(e10)))), inference(forward_demodulation, [], [f4639, f2643])).
fof(f4639, plain, (~ (h18(op1(e10, e12)) = op2(h18(e10), e24)) | ~ (h18(op1(e10, e13)) = op2(h18(e10), h8(e11))) | ~ (h18(op1(e10, e14)) = op2(h18(e10), e21)) | ~ (h18(op1(e11, e10)) = op2(h8(e13), h18(e10))) | ~ (h18(e10) = h18(op1(e11, e11))) | ~ (h18(op1(e11, e12)) = op2(h8(e13), e24)) | ~ (op2(h8(e13), h8(e11)) = h18(op1(e11, e13))) | ~ (h18(op1(e11, e14)) = op2(h8(e13), e21)) | ~ (h18(op1(e12, e10)) = op2(e24, h18(e10))) | ~ (h18(op1(e12, e11)) = op2(e24, h8(e13))) | ~ (op2(e24, e24) = h18(op1(e12, e12))) | ~ (h18(op1(e12, e13)) = op2(e24, h8(e11))) | ~ (h8(e11) = h18(op1(e12, e14))) | ~ (h18(op1(e13, e10)) = op2(h8(e11), h18(e10))) | ~ (op2(h8(e11), h8(e13)) = h18(op1(e13, e11))) | ~ (h18(op1(e13, e12)) = op2(h8(e11), e24)) | ~ (h8(e10) = h18(op1(e13, e13))) | ~ (h18(op1(e13, e14)) = op2(h8(e11), e21)) | ~ (h18(op1(e14, e10)) = op2(e21, h18(e10))) | ~ (h18(op1(e14, e11)) = op2(e21, h8(e13))) | ~ (h8(e13) = h18(op1(e14, e12))) | ~ (h18(op1(e14, e13)) = op2(e21, h8(e11))) | ~ (op2(e21, e21) = h18(op1(e14, e14))) | sP71 | sP70 | sP68 | ~ (h18(op1(e10, e11)) = op2(h18(e10), h18(e11))) | ~ (h18(op1(e10, e10)) = op2(h18(e10), h18(e10)))), inference(forward_demodulation, [], [f4638, f758])).
fof(f4638, plain, (~ (h18(op1(e10, e13)) = op2(h18(e10), h8(e11))) | ~ (h18(op1(e10, e14)) = op2(h18(e10), e21)) | ~ (h18(op1(e11, e10)) = op2(h8(e13), h18(e10))) | ~ (h18(e10) = h18(op1(e11, e11))) | ~ (h18(op1(e11, e12)) = op2(h8(e13), e24)) | ~ (op2(h8(e13), h8(e11)) = h18(op1(e11, e13))) | ~ (h18(op1(e11, e14)) = op2(h8(e13), e21)) | ~ (h18(op1(e12, e10)) = op2(e24, h18(e10))) | ~ (h18(op1(e12, e11)) = op2(e24, h8(e13))) | ~ (op2(e24, e24) = h18(op1(e12, e12))) | ~ (h18(op1(e12, e13)) = op2(e24, h8(e11))) | ~ (h8(e11) = h18(op1(e12, e14))) | ~ (h18(op1(e13, e10)) = op2(h8(e11), h18(e10))) | ~ (op2(h8(e11), h8(e13)) = h18(op1(e13, e11))) | ~ (h18(op1(e13, e12)) = op2(h8(e11), e24)) | ~ (h8(e10) = h18(op1(e13, e13))) | ~ (h18(op1(e13, e14)) = op2(h8(e11), e21)) | ~ (h18(op1(e14, e10)) = op2(e21, h18(e10))) | ~ (h18(op1(e14, e11)) = op2(e21, h8(e13))) | ~ (h8(e13) = h18(op1(e14, e12))) | ~ (h18(op1(e14, e13)) = op2(e21, h8(e11))) | ~ (op2(e21, e21) = h18(op1(e14, e14))) | sP71 | sP70 | sP68 | ~ (h18(op1(e10, e12)) = op2(h18(e10), h18(e12))) | ~ (h18(op1(e10, e11)) = op2(h18(e10), h18(e11))) | ~ (h18(op1(e10, e10)) = op2(h18(e10), h18(e10)))), inference(forward_demodulation, [], [f4637, f2642])).
fof(f4637, plain, (~ (h18(op1(e10, e14)) = op2(h18(e10), e21)) | ~ (h18(op1(e11, e10)) = op2(h8(e13), h18(e10))) | ~ (h18(e10) = h18(op1(e11, e11))) | ~ (h18(op1(e11, e12)) = op2(h8(e13), e24)) | ~ (op2(h8(e13), h8(e11)) = h18(op1(e11, e13))) | ~ (h18(op1(e11, e14)) = op2(h8(e13), e21)) | ~ (h18(op1(e12, e10)) = op2(e24, h18(e10))) | ~ (h18(op1(e12, e11)) = op2(e24, h8(e13))) | ~ (op2(e24, e24) = h18(op1(e12, e12))) | ~ (h18(op1(e12, e13)) = op2(e24, h8(e11))) | ~ (h8(e11) = h18(op1(e12, e14))) | ~ (h18(op1(e13, e10)) = op2(h8(e11), h18(e10))) | ~ (op2(h8(e11), h8(e13)) = h18(op1(e13, e11))) | ~ (h18(op1(e13, e12)) = op2(h8(e11), e24)) | ~ (h8(e10) = h18(op1(e13, e13))) | ~ (h18(op1(e13, e14)) = op2(h8(e11), e21)) | ~ (h18(op1(e14, e10)) = op2(e21, h18(e10))) | ~ (h18(op1(e14, e11)) = op2(e21, h8(e13))) | ~ (h8(e13) = h18(op1(e14, e12))) | ~ (h18(op1(e14, e13)) = op2(e21, h8(e11))) | ~ (op2(e21, e21) = h18(op1(e14, e14))) | sP71 | sP70 | sP68 | ~ (h18(op1(e10, e13)) = op2(h18(e10), h18(e13))) | ~ (h18(op1(e10, e12)) = op2(h18(e10), h18(e12))) | ~ (h18(op1(e10, e11)) = op2(h18(e10), h18(e11))) | ~ (h18(op1(e10, e10)) = op2(h18(e10), h18(e10)))), inference(forward_demodulation, [], [f4636, f759])).
fof(f4636, plain, (~ (h18(op1(e11, e10)) = op2(h8(e13), h18(e10))) | ~ (h18(e10) = h18(op1(e11, e11))) | ~ (h18(op1(e11, e12)) = op2(h8(e13), e24)) | ~ (op2(h8(e13), h8(e11)) = h18(op1(e11, e13))) | ~ (h18(op1(e11, e14)) = op2(h8(e13), e21)) | ~ (h18(op1(e12, e10)) = op2(e24, h18(e10))) | ~ (h18(op1(e12, e11)) = op2(e24, h8(e13))) | ~ (op2(e24, e24) = h18(op1(e12, e12))) | ~ (h18(op1(e12, e13)) = op2(e24, h8(e11))) | ~ (h8(e11) = h18(op1(e12, e14))) | ~ (h18(op1(e13, e10)) = op2(h8(e11), h18(e10))) | ~ (op2(h8(e11), h8(e13)) = h18(op1(e13, e11))) | ~ (h18(op1(e13, e12)) = op2(h8(e11), e24)) | ~ (h8(e10) = h18(op1(e13, e13))) | ~ (h18(op1(e13, e14)) = op2(h8(e11), e21)) | ~ (h18(op1(e14, e10)) = op2(e21, h18(e10))) | ~ (h18(op1(e14, e11)) = op2(e21, h8(e13))) | ~ (h8(e13) = h18(op1(e14, e12))) | ~ (h18(op1(e14, e13)) = op2(e21, h8(e11))) | ~ (op2(e21, e21) = h18(op1(e14, e14))) | sP71 | sP70 | sP68 | ~ (h18(op1(e10, e14)) = op2(h18(e10), h18(e14))) | ~ (h18(op1(e10, e13)) = op2(h18(e10), h18(e13))) | ~ (h18(op1(e10, e12)) = op2(h18(e10), h18(e12))) | ~ (h18(op1(e10, e11)) = op2(h18(e10), h18(e11))) | ~ (h18(op1(e10, e10)) = op2(h18(e10), h18(e10)))), inference(forward_demodulation, [], [f4635, f2643])).
fof(f4635, plain, (~ (h18(e10) = h18(op1(e11, e11))) | ~ (h18(op1(e11, e12)) = op2(h8(e13), e24)) | ~ (op2(h8(e13), h8(e11)) = h18(op1(e11, e13))) | ~ (h18(op1(e11, e14)) = op2(h8(e13), e21)) | ~ (h18(op1(e12, e10)) = op2(e24, h18(e10))) | ~ (h18(op1(e12, e11)) = op2(e24, h8(e13))) | ~ (op2(e24, e24) = h18(op1(e12, e12))) | ~ (h18(op1(e12, e13)) = op2(e24, h8(e11))) | ~ (h8(e11) = h18(op1(e12, e14))) | ~ (h18(op1(e13, e10)) = op2(h8(e11), h18(e10))) | ~ (op2(h8(e11), h8(e13)) = h18(op1(e13, e11))) | ~ (h18(op1(e13, e12)) = op2(h8(e11), e24)) | ~ (h8(e10) = h18(op1(e13, e13))) | ~ (h18(op1(e13, e14)) = op2(h8(e11), e21)) | ~ (h18(op1(e14, e10)) = op2(e21, h18(e10))) | ~ (h18(op1(e14, e11)) = op2(e21, h8(e13))) | ~ (h8(e13) = h18(op1(e14, e12))) | ~ (h18(op1(e14, e13)) = op2(e21, h8(e11))) | ~ (op2(e21, e21) = h18(op1(e14, e14))) | sP71 | sP70 | sP68 | ~ (h18(op1(e11, e10)) = op2(h18(e11), h18(e10))) | ~ (h18(op1(e10, e14)) = op2(h18(e10), h18(e14))) | ~ (h18(op1(e10, e13)) = op2(h18(e10), h18(e13))) | ~ (h18(op1(e10, e12)) = op2(h18(e10), h18(e12))) | ~ (h18(op1(e10, e11)) = op2(h18(e10), h18(e11))) | ~ (h18(op1(e10, e10)) = op2(h18(e10), h18(e10)))), inference(forward_demodulation, [], [f4634, f2644])).
fof(f4634, plain, (~ (op2(h8(e13), h8(e13)) = h18(op1(e11, e11))) | ~ (h18(op1(e11, e12)) = op2(h8(e13), e24)) | ~ (op2(h8(e13), h8(e11)) = h18(op1(e11, e13))) | ~ (h18(op1(e11, e14)) = op2(h8(e13), e21)) | ~ (h18(op1(e12, e10)) = op2(e24, h18(e10))) | ~ (h18(op1(e12, e11)) = op2(e24, h8(e13))) | ~ (op2(e24, e24) = h18(op1(e12, e12))) | ~ (h18(op1(e12, e13)) = op2(e24, h8(e11))) | ~ (h8(e11) = h18(op1(e12, e14))) | ~ (h18(op1(e13, e10)) = op2(h8(e11), h18(e10))) | ~ (op2(h8(e11), h8(e13)) = h18(op1(e13, e11))) | ~ (h18(op1(e13, e12)) = op2(h8(e11), e24)) | ~ (h8(e10) = h18(op1(e13, e13))) | ~ (h18(op1(e13, e14)) = op2(h8(e11), e21)) | ~ (h18(op1(e14, e10)) = op2(e21, h18(e10))) | ~ (h18(op1(e14, e11)) = op2(e21, h8(e13))) | ~ (h8(e13) = h18(op1(e14, e12))) | ~ (h18(op1(e14, e13)) = op2(e21, h8(e11))) | ~ (op2(e21, e21) = h18(op1(e14, e14))) | sP71 | sP70 | sP68 | ~ (h18(op1(e11, e10)) = op2(h18(e11), h18(e10))) | ~ (h18(op1(e10, e14)) = op2(h18(e10), h18(e14))) | ~ (h18(op1(e10, e13)) = op2(h18(e10), h18(e13))) | ~ (h18(op1(e10, e12)) = op2(h18(e10), h18(e12))) | ~ (h18(op1(e10, e11)) = op2(h18(e10), h18(e11))) | ~ (h18(op1(e10, e10)) = op2(h18(e10), h18(e10)))), inference(forward_demodulation, [], [f4633, f2643])).
fof(f4633, plain, (~ (h18(op1(e11, e12)) = op2(h8(e13), e24)) | ~ (op2(h8(e13), h8(e11)) = h18(op1(e11, e13))) | ~ (h18(op1(e11, e14)) = op2(h8(e13), e21)) | ~ (h18(op1(e12, e10)) = op2(e24, h18(e10))) | ~ (h18(op1(e12, e11)) = op2(e24, h8(e13))) | ~ (op2(e24, e24) = h18(op1(e12, e12))) | ~ (h18(op1(e12, e13)) = op2(e24, h8(e11))) | ~ (h8(e11) = h18(op1(e12, e14))) | ~ (h18(op1(e13, e10)) = op2(h8(e11), h18(e10))) | ~ (op2(h8(e11), h8(e13)) = h18(op1(e13, e11))) | ~ (h18(op1(e13, e12)) = op2(h8(e11), e24)) | ~ (h8(e10) = h18(op1(e13, e13))) | ~ (h18(op1(e13, e14)) = op2(h8(e11), e21)) | ~ (h18(op1(e14, e10)) = op2(e21, h18(e10))) | ~ (h18(op1(e14, e11)) = op2(e21, h8(e13))) | ~ (h8(e13) = h18(op1(e14, e12))) | ~ (h18(op1(e14, e13)) = op2(e21, h8(e11))) | ~ (op2(e21, e21) = h18(op1(e14, e14))) | sP71 | sP70 | sP68 | ~ (h18(op1(e11, e11)) = op2(h18(e11), h18(e11))) | ~ (h18(op1(e11, e10)) = op2(h18(e11), h18(e10))) | ~ (h18(op1(e10, e14)) = op2(h18(e10), h18(e14))) | ~ (h18(op1(e10, e13)) = op2(h18(e10), h18(e13))) | ~ (h18(op1(e10, e12)) = op2(h18(e10), h18(e12))) | ~ (h18(op1(e10, e11)) = op2(h18(e10), h18(e11))) | ~ (h18(op1(e10, e10)) = op2(h18(e10), h18(e10)))), inference(forward_demodulation, [], [f4632, f2643])).
fof(f4632, plain, (~ (h18(op1(e11, e12)) = op2(h18(e11), e24)) | ~ (op2(h8(e13), h8(e11)) = h18(op1(e11, e13))) | ~ (h18(op1(e11, e14)) = op2(h8(e13), e21)) | ~ (h18(op1(e12, e10)) = op2(e24, h18(e10))) | ~ (h18(op1(e12, e11)) = op2(e24, h8(e13))) | ~ (op2(e24, e24) = h18(op1(e12, e12))) | ~ (h18(op1(e12, e13)) = op2(e24, h8(e11))) | ~ (h8(e11) = h18(op1(e12, e14))) | ~ (h18(op1(e13, e10)) = op2(h8(e11), h18(e10))) | ~ (op2(h8(e11), h8(e13)) = h18(op1(e13, e11))) | ~ (h18(op1(e13, e12)) = op2(h8(e11), e24)) | ~ (h8(e10) = h18(op1(e13, e13))) | ~ (h18(op1(e13, e14)) = op2(h8(e11), e21)) | ~ (h18(op1(e14, e10)) = op2(e21, h18(e10))) | ~ (h18(op1(e14, e11)) = op2(e21, h8(e13))) | ~ (h8(e13) = h18(op1(e14, e12))) | ~ (h18(op1(e14, e13)) = op2(e21, h8(e11))) | ~ (op2(e21, e21) = h18(op1(e14, e14))) | sP71 | sP70 | sP68 | ~ (h18(op1(e11, e11)) = op2(h18(e11), h18(e11))) | ~ (h18(op1(e11, e10)) = op2(h18(e11), h18(e10))) | ~ (h18(op1(e10, e14)) = op2(h18(e10), h18(e14))) | ~ (h18(op1(e10, e13)) = op2(h18(e10), h18(e13))) | ~ (h18(op1(e10, e12)) = op2(h18(e10), h18(e12))) | ~ (h18(op1(e10, e11)) = op2(h18(e10), h18(e11))) | ~ (h18(op1(e10, e10)) = op2(h18(e10), h18(e10)))), inference(forward_demodulation, [], [f4631, f758])).
fof(f4631, plain, (~ (op2(h8(e13), h8(e11)) = h18(op1(e11, e13))) | ~ (h18(op1(e11, e14)) = op2(h8(e13), e21)) | ~ (h18(op1(e12, e10)) = op2(e24, h18(e10))) | ~ (h18(op1(e12, e11)) = op2(e24, h8(e13))) | ~ (op2(e24, e24) = h18(op1(e12, e12))) | ~ (h18(op1(e12, e13)) = op2(e24, h8(e11))) | ~ (h8(e11) = h18(op1(e12, e14))) | ~ (h18(op1(e13, e10)) = op2(h8(e11), h18(e10))) | ~ (op2(h8(e11), h8(e13)) = h18(op1(e13, e11))) | ~ (h18(op1(e13, e12)) = op2(h8(e11), e24)) | ~ (h8(e10) = h18(op1(e13, e13))) | ~ (h18(op1(e13, e14)) = op2(h8(e11), e21)) | ~ (h18(op1(e14, e10)) = op2(e21, h18(e10))) | ~ (h18(op1(e14, e11)) = op2(e21, h8(e13))) | ~ (h8(e13) = h18(op1(e14, e12))) | ~ (h18(op1(e14, e13)) = op2(e21, h8(e11))) | ~ (op2(e21, e21) = h18(op1(e14, e14))) | sP71 | sP70 | sP68 | ~ (h18(op1(e11, e12)) = op2(h18(e11), h18(e12))) | ~ (h18(op1(e11, e11)) = op2(h18(e11), h18(e11))) | ~ (h18(op1(e11, e10)) = op2(h18(e11), h18(e10))) | ~ (h18(op1(e10, e14)) = op2(h18(e10), h18(e14))) | ~ (h18(op1(e10, e13)) = op2(h18(e10), h18(e13))) | ~ (h18(op1(e10, e12)) = op2(h18(e10), h18(e12))) | ~ (h18(op1(e10, e11)) = op2(h18(e10), h18(e11))) | ~ (h18(op1(e10, e10)) = op2(h18(e10), h18(e10)))), inference(forward_demodulation, [], [f4630, f2643])).
fof(f4630, plain, (~ (h18(op1(e11, e13)) = op2(h18(e11), h8(e11))) | ~ (h18(op1(e11, e14)) = op2(h8(e13), e21)) | ~ (h18(op1(e12, e10)) = op2(e24, h18(e10))) | ~ (h18(op1(e12, e11)) = op2(e24, h8(e13))) | ~ (op2(e24, e24) = h18(op1(e12, e12))) | ~ (h18(op1(e12, e13)) = op2(e24, h8(e11))) | ~ (h8(e11) = h18(op1(e12, e14))) | ~ (h18(op1(e13, e10)) = op2(h8(e11), h18(e10))) | ~ (op2(h8(e11), h8(e13)) = h18(op1(e13, e11))) | ~ (h18(op1(e13, e12)) = op2(h8(e11), e24)) | ~ (h8(e10) = h18(op1(e13, e13))) | ~ (h18(op1(e13, e14)) = op2(h8(e11), e21)) | ~ (h18(op1(e14, e10)) = op2(e21, h18(e10))) | ~ (h18(op1(e14, e11)) = op2(e21, h8(e13))) | ~ (h8(e13) = h18(op1(e14, e12))) | ~ (h18(op1(e14, e13)) = op2(e21, h8(e11))) | ~ (op2(e21, e21) = h18(op1(e14, e14))) | sP71 | sP70 | sP68 | ~ (h18(op1(e11, e12)) = op2(h18(e11), h18(e12))) | ~ (h18(op1(e11, e11)) = op2(h18(e11), h18(e11))) | ~ (h18(op1(e11, e10)) = op2(h18(e11), h18(e10))) | ~ (h18(op1(e10, e14)) = op2(h18(e10), h18(e14))) | ~ (h18(op1(e10, e13)) = op2(h18(e10), h18(e13))) | ~ (h18(op1(e10, e12)) = op2(h18(e10), h18(e12))) | ~ (h18(op1(e10, e11)) = op2(h18(e10), h18(e11))) | ~ (h18(op1(e10, e10)) = op2(h18(e10), h18(e10)))), inference(forward_demodulation, [], [f4629, f2642])).
fof(f4629, plain, (~ (h18(op1(e11, e14)) = op2(h8(e13), e21)) | ~ (h18(op1(e12, e10)) = op2(e24, h18(e10))) | ~ (h18(op1(e12, e11)) = op2(e24, h8(e13))) | ~ (op2(e24, e24) = h18(op1(e12, e12))) | ~ (h18(op1(e12, e13)) = op2(e24, h8(e11))) | ~ (h8(e11) = h18(op1(e12, e14))) | ~ (h18(op1(e13, e10)) = op2(h8(e11), h18(e10))) | ~ (op2(h8(e11), h8(e13)) = h18(op1(e13, e11))) | ~ (h18(op1(e13, e12)) = op2(h8(e11), e24)) | ~ (h8(e10) = h18(op1(e13, e13))) | ~ (h18(op1(e13, e14)) = op2(h8(e11), e21)) | ~ (h18(op1(e14, e10)) = op2(e21, h18(e10))) | ~ (h18(op1(e14, e11)) = op2(e21, h8(e13))) | ~ (h8(e13) = h18(op1(e14, e12))) | ~ (h18(op1(e14, e13)) = op2(e21, h8(e11))) | ~ (op2(e21, e21) = h18(op1(e14, e14))) | sP71 | sP70 | sP68 | ~ (h18(op1(e11, e13)) = op2(h18(e11), h18(e13))) | ~ (h18(op1(e11, e12)) = op2(h18(e11), h18(e12))) | ~ (h18(op1(e11, e11)) = op2(h18(e11), h18(e11))) | ~ (h18(op1(e11, e10)) = op2(h18(e11), h18(e10))) | ~ (h18(op1(e10, e14)) = op2(h18(e10), h18(e14))) | ~ (h18(op1(e10, e13)) = op2(h18(e10), h18(e13))) | ~ (h18(op1(e10, e12)) = op2(h18(e10), h18(e12))) | ~ (h18(op1(e10, e11)) = op2(h18(e10), h18(e11))) | ~ (h18(op1(e10, e10)) = op2(h18(e10), h18(e10)))), inference(forward_demodulation, [], [f4628, f2643])).
fof(f4628, plain, (~ (h18(op1(e11, e14)) = op2(h18(e11), e21)) | ~ (h18(op1(e12, e10)) = op2(e24, h18(e10))) | ~ (h18(op1(e12, e11)) = op2(e24, h8(e13))) | ~ (op2(e24, e24) = h18(op1(e12, e12))) | ~ (h18(op1(e12, e13)) = op2(e24, h8(e11))) | ~ (h8(e11) = h18(op1(e12, e14))) | ~ (h18(op1(e13, e10)) = op2(h8(e11), h18(e10))) | ~ (op2(h8(e11), h8(e13)) = h18(op1(e13, e11))) | ~ (h18(op1(e13, e12)) = op2(h8(e11), e24)) | ~ (h8(e10) = h18(op1(e13, e13))) | ~ (h18(op1(e13, e14)) = op2(h8(e11), e21)) | ~ (h18(op1(e14, e10)) = op2(e21, h18(e10))) | ~ (h18(op1(e14, e11)) = op2(e21, h8(e13))) | ~ (h8(e13) = h18(op1(e14, e12))) | ~ (h18(op1(e14, e13)) = op2(e21, h8(e11))) | ~ (op2(e21, e21) = h18(op1(e14, e14))) | sP71 | sP70 | sP68 | ~ (h18(op1(e11, e13)) = op2(h18(e11), h18(e13))) | ~ (h18(op1(e11, e12)) = op2(h18(e11), h18(e12))) | ~ (h18(op1(e11, e11)) = op2(h18(e11), h18(e11))) | ~ (h18(op1(e11, e10)) = op2(h18(e11), h18(e10))) | ~ (h18(op1(e10, e14)) = op2(h18(e10), h18(e14))) | ~ (h18(op1(e10, e13)) = op2(h18(e10), h18(e13))) | ~ (h18(op1(e10, e12)) = op2(h18(e10), h18(e12))) | ~ (h18(op1(e10, e11)) = op2(h18(e10), h18(e11))) | ~ (h18(op1(e10, e10)) = op2(h18(e10), h18(e10)))), inference(forward_demodulation, [], [f4627, f759])).
fof(f4627, plain, (~ (h18(op1(e12, e10)) = op2(e24, h18(e10))) | ~ (h18(op1(e12, e11)) = op2(e24, h8(e13))) | ~ (op2(e24, e24) = h18(op1(e12, e12))) | ~ (h18(op1(e12, e13)) = op2(e24, h8(e11))) | ~ (h8(e11) = h18(op1(e12, e14))) | ~ (h18(op1(e13, e10)) = op2(h8(e11), h18(e10))) | ~ (op2(h8(e11), h8(e13)) = h18(op1(e13, e11))) | ~ (h18(op1(e13, e12)) = op2(h8(e11), e24)) | ~ (h8(e10) = h18(op1(e13, e13))) | ~ (h18(op1(e13, e14)) = op2(h8(e11), e21)) | ~ (h18(op1(e14, e10)) = op2(e21, h18(e10))) | ~ (h18(op1(e14, e11)) = op2(e21, h8(e13))) | ~ (h8(e13) = h18(op1(e14, e12))) | ~ (h18(op1(e14, e13)) = op2(e21, h8(e11))) | ~ (op2(e21, e21) = h18(op1(e14, e14))) | sP71 | sP70 | sP68 | ~ (h18(op1(e11, e14)) = op2(h18(e11), h18(e14))) | ~ (h18(op1(e11, e13)) = op2(h18(e11), h18(e13))) | ~ (h18(op1(e11, e12)) = op2(h18(e11), h18(e12))) | ~ (h18(op1(e11, e11)) = op2(h18(e11), h18(e11))) | ~ (h18(op1(e11, e10)) = op2(h18(e11), h18(e10))) | ~ (h18(op1(e10, e14)) = op2(h18(e10), h18(e14))) | ~ (h18(op1(e10, e13)) = op2(h18(e10), h18(e13))) | ~ (h18(op1(e10, e12)) = op2(h18(e10), h18(e12))) | ~ (h18(op1(e10, e11)) = op2(h18(e10), h18(e11))) | ~ (h18(op1(e10, e10)) = op2(h18(e10), h18(e10)))), inference(forward_demodulation, [], [f4626, f758])).
fof(f4626, plain, (~ (h18(op1(e12, e11)) = op2(e24, h8(e13))) | ~ (op2(e24, e24) = h18(op1(e12, e12))) | ~ (h18(op1(e12, e13)) = op2(e24, h8(e11))) | ~ (h8(e11) = h18(op1(e12, e14))) | ~ (h18(op1(e13, e10)) = op2(h8(e11), h18(e10))) | ~ (op2(h8(e11), h8(e13)) = h18(op1(e13, e11))) | ~ (h18(op1(e13, e12)) = op2(h8(e11), e24)) | ~ (h8(e10) = h18(op1(e13, e13))) | ~ (h18(op1(e13, e14)) = op2(h8(e11), e21)) | ~ (h18(op1(e14, e10)) = op2(e21, h18(e10))) | ~ (h18(op1(e14, e11)) = op2(e21, h8(e13))) | ~ (h8(e13) = h18(op1(e14, e12))) | ~ (h18(op1(e14, e13)) = op2(e21, h8(e11))) | ~ (op2(e21, e21) = h18(op1(e14, e14))) | sP71 | sP70 | sP68 | ~ (h18(op1(e12, e10)) = op2(h18(e12), h18(e10))) | ~ (h18(op1(e11, e14)) = op2(h18(e11), h18(e14))) | ~ (h18(op1(e11, e13)) = op2(h18(e11), h18(e13))) | ~ (h18(op1(e11, e12)) = op2(h18(e11), h18(e12))) | ~ (h18(op1(e11, e11)) = op2(h18(e11), h18(e11))) | ~ (h18(op1(e11, e10)) = op2(h18(e11), h18(e10))) | ~ (h18(op1(e10, e14)) = op2(h18(e10), h18(e14))) | ~ (h18(op1(e10, e13)) = op2(h18(e10), h18(e13))) | ~ (h18(op1(e10, e12)) = op2(h18(e10), h18(e12))) | ~ (h18(op1(e10, e11)) = op2(h18(e10), h18(e11))) | ~ (h18(op1(e10, e10)) = op2(h18(e10), h18(e10)))), inference(forward_demodulation, [], [f4625, f758])).
fof(f4625, plain, (~ (h18(op1(e12, e11)) = op2(h18(e12), h8(e13))) | ~ (op2(e24, e24) = h18(op1(e12, e12))) | ~ (h18(op1(e12, e13)) = op2(e24, h8(e11))) | ~ (h8(e11) = h18(op1(e12, e14))) | ~ (h18(op1(e13, e10)) = op2(h8(e11), h18(e10))) | ~ (op2(h8(e11), h8(e13)) = h18(op1(e13, e11))) | ~ (h18(op1(e13, e12)) = op2(h8(e11), e24)) | ~ (h8(e10) = h18(op1(e13, e13))) | ~ (h18(op1(e13, e14)) = op2(h8(e11), e21)) | ~ (h18(op1(e14, e10)) = op2(e21, h18(e10))) | ~ (h18(op1(e14, e11)) = op2(e21, h8(e13))) | ~ (h8(e13) = h18(op1(e14, e12))) | ~ (h18(op1(e14, e13)) = op2(e21, h8(e11))) | ~ (op2(e21, e21) = h18(op1(e14, e14))) | sP71 | sP70 | sP68 | ~ (h18(op1(e12, e10)) = op2(h18(e12), h18(e10))) | ~ (h18(op1(e11, e14)) = op2(h18(e11), h18(e14))) | ~ (h18(op1(e11, e13)) = op2(h18(e11), h18(e13))) | ~ (h18(op1(e11, e12)) = op2(h18(e11), h18(e12))) | ~ (h18(op1(e11, e11)) = op2(h18(e11), h18(e11))) | ~ (h18(op1(e11, e10)) = op2(h18(e11), h18(e10))) | ~ (h18(op1(e10, e14)) = op2(h18(e10), h18(e14))) | ~ (h18(op1(e10, e13)) = op2(h18(e10), h18(e13))) | ~ (h18(op1(e10, e12)) = op2(h18(e10), h18(e12))) | ~ (h18(op1(e10, e11)) = op2(h18(e10), h18(e11))) | ~ (h18(op1(e10, e10)) = op2(h18(e10), h18(e10)))), inference(forward_demodulation, [], [f4624, f2643])).
fof(f4624, plain, (~ (op2(e24, e24) = h18(op1(e12, e12))) | ~ (h18(op1(e12, e13)) = op2(e24, h8(e11))) | ~ (h8(e11) = h18(op1(e12, e14))) | ~ (h18(op1(e13, e10)) = op2(h8(e11), h18(e10))) | ~ (op2(h8(e11), h8(e13)) = h18(op1(e13, e11))) | ~ (h18(op1(e13, e12)) = op2(h8(e11), e24)) | ~ (h8(e10) = h18(op1(e13, e13))) | ~ (h18(op1(e13, e14)) = op2(h8(e11), e21)) | ~ (h18(op1(e14, e10)) = op2(e21, h18(e10))) | ~ (h18(op1(e14, e11)) = op2(e21, h8(e13))) | ~ (h8(e13) = h18(op1(e14, e12))) | ~ (h18(op1(e14, e13)) = op2(e21, h8(e11))) | ~ (op2(e21, e21) = h18(op1(e14, e14))) | sP71 | sP70 | sP68 | ~ (h18(op1(e12, e11)) = op2(h18(e12), h18(e11))) | ~ (h18(op1(e12, e10)) = op2(h18(e12), h18(e10))) | ~ (h18(op1(e11, e14)) = op2(h18(e11), h18(e14))) | ~ (h18(op1(e11, e13)) = op2(h18(e11), h18(e13))) | ~ (h18(op1(e11, e12)) = op2(h18(e11), h18(e12))) | ~ (h18(op1(e11, e11)) = op2(h18(e11), h18(e11))) | ~ (h18(op1(e11, e10)) = op2(h18(e11), h18(e10))) | ~ (h18(op1(e10, e14)) = op2(h18(e10), h18(e14))) | ~ (h18(op1(e10, e13)) = op2(h18(e10), h18(e13))) | ~ (h18(op1(e10, e12)) = op2(h18(e10), h18(e12))) | ~ (h18(op1(e10, e11)) = op2(h18(e10), h18(e11))) | ~ (h18(op1(e10, e10)) = op2(h18(e10), h18(e10)))), inference(forward_demodulation, [], [f4623, f758])).
fof(f4623, plain, (~ (h18(op1(e12, e13)) = op2(e24, h8(e11))) | ~ (h8(e11) = h18(op1(e12, e14))) | ~ (h18(op1(e13, e10)) = op2(h8(e11), h18(e10))) | ~ (op2(h8(e11), h8(e13)) = h18(op1(e13, e11))) | ~ (h18(op1(e13, e12)) = op2(h8(e11), e24)) | ~ (h8(e10) = h18(op1(e13, e13))) | ~ (h18(op1(e13, e14)) = op2(h8(e11), e21)) | ~ (h18(op1(e14, e10)) = op2(e21, h18(e10))) | ~ (h18(op1(e14, e11)) = op2(e21, h8(e13))) | ~ (h8(e13) = h18(op1(e14, e12))) | ~ (h18(op1(e14, e13)) = op2(e21, h8(e11))) | ~ (op2(e21, e21) = h18(op1(e14, e14))) | sP71 | sP70 | sP68 | ~ (h18(op1(e12, e12)) = op2(h18(e12), h18(e12))) | ~ (h18(op1(e12, e11)) = op2(h18(e12), h18(e11))) | ~ (h18(op1(e12, e10)) = op2(h18(e12), h18(e10))) | ~ (h18(op1(e11, e14)) = op2(h18(e11), h18(e14))) | ~ (h18(op1(e11, e13)) = op2(h18(e11), h18(e13))) | ~ (h18(op1(e11, e12)) = op2(h18(e11), h18(e12))) | ~ (h18(op1(e11, e11)) = op2(h18(e11), h18(e11))) | ~ (h18(op1(e11, e10)) = op2(h18(e11), h18(e10))) | ~ (h18(op1(e10, e14)) = op2(h18(e10), h18(e14))) | ~ (h18(op1(e10, e13)) = op2(h18(e10), h18(e13))) | ~ (h18(op1(e10, e12)) = op2(h18(e10), h18(e12))) | ~ (h18(op1(e10, e11)) = op2(h18(e10), h18(e11))) | ~ (h18(op1(e10, e10)) = op2(h18(e10), h18(e10)))), inference(forward_demodulation, [], [f4622, f758])).
fof(f4622, plain, (~ (h18(op1(e12, e13)) = op2(h18(e12), h8(e11))) | ~ (h8(e11) = h18(op1(e12, e14))) | ~ (h18(op1(e13, e10)) = op2(h8(e11), h18(e10))) | ~ (op2(h8(e11), h8(e13)) = h18(op1(e13, e11))) | ~ (h18(op1(e13, e12)) = op2(h8(e11), e24)) | ~ (h8(e10) = h18(op1(e13, e13))) | ~ (h18(op1(e13, e14)) = op2(h8(e11), e21)) | ~ (h18(op1(e14, e10)) = op2(e21, h18(e10))) | ~ (h18(op1(e14, e11)) = op2(e21, h8(e13))) | ~ (h8(e13) = h18(op1(e14, e12))) | ~ (h18(op1(e14, e13)) = op2(e21, h8(e11))) | ~ (op2(e21, e21) = h18(op1(e14, e14))) | sP71 | sP70 | sP68 | ~ (h18(op1(e12, e12)) = op2(h18(e12), h18(e12))) | ~ (h18(op1(e12, e11)) = op2(h18(e12), h18(e11))) | ~ (h18(op1(e12, e10)) = op2(h18(e12), h18(e10))) | ~ (h18(op1(e11, e14)) = op2(h18(e11), h18(e14))) | ~ (h18(op1(e11, e13)) = op2(h18(e11), h18(e13))) | ~ (h18(op1(e11, e12)) = op2(h18(e11), h18(e12))) | ~ (h18(op1(e11, e11)) = op2(h18(e11), h18(e11))) | ~ (h18(op1(e11, e10)) = op2(h18(e11), h18(e10))) | ~ (h18(op1(e10, e14)) = op2(h18(e10), h18(e14))) | ~ (h18(op1(e10, e13)) = op2(h18(e10), h18(e13))) | ~ (h18(op1(e10, e12)) = op2(h18(e10), h18(e12))) | ~ (h18(op1(e10, e11)) = op2(h18(e10), h18(e11))) | ~ (h18(op1(e10, e10)) = op2(h18(e10), h18(e10)))), inference(forward_demodulation, [], [f4621, f2642])).
fof(f4621, plain, (~ (h8(e11) = h18(op1(e12, e14))) | ~ (h18(op1(e13, e10)) = op2(h8(e11), h18(e10))) | ~ (op2(h8(e11), h8(e13)) = h18(op1(e13, e11))) | ~ (h18(op1(e13, e12)) = op2(h8(e11), e24)) | ~ (h8(e10) = h18(op1(e13, e13))) | ~ (h18(op1(e13, e14)) = op2(h8(e11), e21)) | ~ (h18(op1(e14, e10)) = op2(e21, h18(e10))) | ~ (h18(op1(e14, e11)) = op2(e21, h8(e13))) | ~ (h8(e13) = h18(op1(e14, e12))) | ~ (h18(op1(e14, e13)) = op2(e21, h8(e11))) | ~ (op2(e21, e21) = h18(op1(e14, e14))) | sP71 | sP70 | sP68 | ~ (h18(op1(e12, e13)) = op2(h18(e12), h18(e13))) | ~ (h18(op1(e12, e12)) = op2(h18(e12), h18(e12))) | ~ (h18(op1(e12, e11)) = op2(h18(e12), h18(e11))) | ~ (h18(op1(e12, e10)) = op2(h18(e12), h18(e10))) | ~ (h18(op1(e11, e14)) = op2(h18(e11), h18(e14))) | ~ (h18(op1(e11, e13)) = op2(h18(e11), h18(e13))) | ~ (h18(op1(e11, e12)) = op2(h18(e11), h18(e12))) | ~ (h18(op1(e11, e11)) = op2(h18(e11), h18(e11))) | ~ (h18(op1(e11, e10)) = op2(h18(e11), h18(e10))) | ~ (h18(op1(e10, e14)) = op2(h18(e10), h18(e14))) | ~ (h18(op1(e10, e13)) = op2(h18(e10), h18(e13))) | ~ (h18(op1(e10, e12)) = op2(h18(e10), h18(e12))) | ~ (h18(op1(e10, e11)) = op2(h18(e10), h18(e11))) | ~ (h18(op1(e10, e10)) = op2(h18(e10), h18(e10)))), inference(forward_demodulation, [], [f4620, f711])).
fof(f4620, plain, (~ (op2(e24, e21) = h18(op1(e12, e14))) | ~ (h18(op1(e13, e10)) = op2(h8(e11), h18(e10))) | ~ (op2(h8(e11), h8(e13)) = h18(op1(e13, e11))) | ~ (h18(op1(e13, e12)) = op2(h8(e11), e24)) | ~ (h8(e10) = h18(op1(e13, e13))) | ~ (h18(op1(e13, e14)) = op2(h8(e11), e21)) | ~ (h18(op1(e14, e10)) = op2(e21, h18(e10))) | ~ (h18(op1(e14, e11)) = op2(e21, h8(e13))) | ~ (h8(e13) = h18(op1(e14, e12))) | ~ (h18(op1(e14, e13)) = op2(e21, h8(e11))) | ~ (op2(e21, e21) = h18(op1(e14, e14))) | sP71 | sP70 | sP68 | ~ (h18(op1(e12, e13)) = op2(h18(e12), h18(e13))) | ~ (h18(op1(e12, e12)) = op2(h18(e12), h18(e12))) | ~ (h18(op1(e12, e11)) = op2(h18(e12), h18(e11))) | ~ (h18(op1(e12, e10)) = op2(h18(e12), h18(e10))) | ~ (h18(op1(e11, e14)) = op2(h18(e11), h18(e14))) | ~ (h18(op1(e11, e13)) = op2(h18(e11), h18(e13))) | ~ (h18(op1(e11, e12)) = op2(h18(e11), h18(e12))) | ~ (h18(op1(e11, e11)) = op2(h18(e11), h18(e11))) | ~ (h18(op1(e11, e10)) = op2(h18(e11), h18(e10))) | ~ (h18(op1(e10, e14)) = op2(h18(e10), h18(e14))) | ~ (h18(op1(e10, e13)) = op2(h18(e10), h18(e13))) | ~ (h18(op1(e10, e12)) = op2(h18(e10), h18(e12))) | ~ (h18(op1(e10, e11)) = op2(h18(e10), h18(e11))) | ~ (h18(op1(e10, e10)) = op2(h18(e10), h18(e10)))), inference(forward_demodulation, [], [f4619, f758])).
fof(f4619, plain, (~ (h18(op1(e12, e14)) = op2(h18(e12), e21)) | ~ (h18(op1(e13, e10)) = op2(h8(e11), h18(e10))) | ~ (op2(h8(e11), h8(e13)) = h18(op1(e13, e11))) | ~ (h18(op1(e13, e12)) = op2(h8(e11), e24)) | ~ (h8(e10) = h18(op1(e13, e13))) | ~ (h18(op1(e13, e14)) = op2(h8(e11), e21)) | ~ (h18(op1(e14, e10)) = op2(e21, h18(e10))) | ~ (h18(op1(e14, e11)) = op2(e21, h8(e13))) | ~ (h8(e13) = h18(op1(e14, e12))) | ~ (h18(op1(e14, e13)) = op2(e21, h8(e11))) | ~ (op2(e21, e21) = h18(op1(e14, e14))) | sP71 | sP70 | sP68 | ~ (h18(op1(e12, e13)) = op2(h18(e12), h18(e13))) | ~ (h18(op1(e12, e12)) = op2(h18(e12), h18(e12))) | ~ (h18(op1(e12, e11)) = op2(h18(e12), h18(e11))) | ~ (h18(op1(e12, e10)) = op2(h18(e12), h18(e10))) | ~ (h18(op1(e11, e14)) = op2(h18(e11), h18(e14))) | ~ (h18(op1(e11, e13)) = op2(h18(e11), h18(e13))) | ~ (h18(op1(e11, e12)) = op2(h18(e11), h18(e12))) | ~ (h18(op1(e11, e11)) = op2(h18(e11), h18(e11))) | ~ (h18(op1(e11, e10)) = op2(h18(e11), h18(e10))) | ~ (h18(op1(e10, e14)) = op2(h18(e10), h18(e14))) | ~ (h18(op1(e10, e13)) = op2(h18(e10), h18(e13))) | ~ (h18(op1(e10, e12)) = op2(h18(e10), h18(e12))) | ~ (h18(op1(e10, e11)) = op2(h18(e10), h18(e11))) | ~ (h18(op1(e10, e10)) = op2(h18(e10), h18(e10)))), inference(forward_demodulation, [], [f4618, f759])).
fof(f4618, plain, (~ (h18(op1(e13, e10)) = op2(h8(e11), h18(e10))) | ~ (op2(h8(e11), h8(e13)) = h18(op1(e13, e11))) | ~ (h18(op1(e13, e12)) = op2(h8(e11), e24)) | ~ (h8(e10) = h18(op1(e13, e13))) | ~ (h18(op1(e13, e14)) = op2(h8(e11), e21)) | ~ (h18(op1(e14, e10)) = op2(e21, h18(e10))) | ~ (h18(op1(e14, e11)) = op2(e21, h8(e13))) | ~ (h8(e13) = h18(op1(e14, e12))) | ~ (h18(op1(e14, e13)) = op2(e21, h8(e11))) | ~ (op2(e21, e21) = h18(op1(e14, e14))) | sP71 | sP70 | sP68 | ~ (h18(op1(e12, e14)) = op2(h18(e12), h18(e14))) | ~ (h18(op1(e12, e13)) = op2(h18(e12), h18(e13))) | ~ (h18(op1(e12, e12)) = op2(h18(e12), h18(e12))) | ~ (h18(op1(e12, e11)) = op2(h18(e12), h18(e11))) | ~ (h18(op1(e12, e10)) = op2(h18(e12), h18(e10))) | ~ (h18(op1(e11, e14)) = op2(h18(e11), h18(e14))) | ~ (h18(op1(e11, e13)) = op2(h18(e11), h18(e13))) | ~ (h18(op1(e11, e12)) = op2(h18(e11), h18(e12))) | ~ (h18(op1(e11, e11)) = op2(h18(e11), h18(e11))) | ~ (h18(op1(e11, e10)) = op2(h18(e11), h18(e10))) | ~ (h18(op1(e10, e14)) = op2(h18(e10), h18(e14))) | ~ (h18(op1(e10, e13)) = op2(h18(e10), h18(e13))) | ~ (h18(op1(e10, e12)) = op2(h18(e10), h18(e12))) | ~ (h18(op1(e10, e11)) = op2(h18(e10), h18(e11))) | ~ (h18(op1(e10, e10)) = op2(h18(e10), h18(e10)))), inference(forward_demodulation, [], [f4617, f2642])).
fof(f4617, plain, (~ (op2(h8(e11), h8(e13)) = h18(op1(e13, e11))) | ~ (h18(op1(e13, e12)) = op2(h8(e11), e24)) | ~ (h8(e10) = h18(op1(e13, e13))) | ~ (h18(op1(e13, e14)) = op2(h8(e11), e21)) | ~ (h18(op1(e14, e10)) = op2(e21, h18(e10))) | ~ (h18(op1(e14, e11)) = op2(e21, h8(e13))) | ~ (h8(e13) = h18(op1(e14, e12))) | ~ (h18(op1(e14, e13)) = op2(e21, h8(e11))) | ~ (op2(e21, e21) = h18(op1(e14, e14))) | sP71 | sP70 | sP68 | ~ (h18(op1(e13, e10)) = op2(h18(e13), h18(e10))) | ~ (h18(op1(e12, e14)) = op2(h18(e12), h18(e14))) | ~ (h18(op1(e12, e13)) = op2(h18(e12), h18(e13))) | ~ (h18(op1(e12, e12)) = op2(h18(e12), h18(e12))) | ~ (h18(op1(e12, e11)) = op2(h18(e12), h18(e11))) | ~ (h18(op1(e12, e10)) = op2(h18(e12), h18(e10))) | ~ (h18(op1(e11, e14)) = op2(h18(e11), h18(e14))) | ~ (h18(op1(e11, e13)) = op2(h18(e11), h18(e13))) | ~ (h18(op1(e11, e12)) = op2(h18(e11), h18(e12))) | ~ (h18(op1(e11, e11)) = op2(h18(e11), h18(e11))) | ~ (h18(op1(e11, e10)) = op2(h18(e11), h18(e10))) | ~ (h18(op1(e10, e14)) = op2(h18(e10), h18(e14))) | ~ (h18(op1(e10, e13)) = op2(h18(e10), h18(e13))) | ~ (h18(op1(e10, e12)) = op2(h18(e10), h18(e12))) | ~ (h18(op1(e10, e11)) = op2(h18(e10), h18(e11))) | ~ (h18(op1(e10, e10)) = op2(h18(e10), h18(e10)))), inference(forward_demodulation, [], [f4616, f2642])).
fof(f4616, plain, (~ (h18(op1(e13, e11)) = op2(h18(e13), h8(e13))) | ~ (h18(op1(e13, e12)) = op2(h8(e11), e24)) | ~ (h8(e10) = h18(op1(e13, e13))) | ~ (h18(op1(e13, e14)) = op2(h8(e11), e21)) | ~ (h18(op1(e14, e10)) = op2(e21, h18(e10))) | ~ (h18(op1(e14, e11)) = op2(e21, h8(e13))) | ~ (h8(e13) = h18(op1(e14, e12))) | ~ (h18(op1(e14, e13)) = op2(e21, h8(e11))) | ~ (op2(e21, e21) = h18(op1(e14, e14))) | sP71 | sP70 | sP68 | ~ (h18(op1(e13, e10)) = op2(h18(e13), h18(e10))) | ~ (h18(op1(e12, e14)) = op2(h18(e12), h18(e14))) | ~ (h18(op1(e12, e13)) = op2(h18(e12), h18(e13))) | ~ (h18(op1(e12, e12)) = op2(h18(e12), h18(e12))) | ~ (h18(op1(e12, e11)) = op2(h18(e12), h18(e11))) | ~ (h18(op1(e12, e10)) = op2(h18(e12), h18(e10))) | ~ (h18(op1(e11, e14)) = op2(h18(e11), h18(e14))) | ~ (h18(op1(e11, e13)) = op2(h18(e11), h18(e13))) | ~ (h18(op1(e11, e12)) = op2(h18(e11), h18(e12))) | ~ (h18(op1(e11, e11)) = op2(h18(e11), h18(e11))) | ~ (h18(op1(e11, e10)) = op2(h18(e11), h18(e10))) | ~ (h18(op1(e10, e14)) = op2(h18(e10), h18(e14))) | ~ (h18(op1(e10, e13)) = op2(h18(e10), h18(e13))) | ~ (h18(op1(e10, e12)) = op2(h18(e10), h18(e12))) | ~ (h18(op1(e10, e11)) = op2(h18(e10), h18(e11))) | ~ (h18(op1(e10, e10)) = op2(h18(e10), h18(e10)))), inference(forward_demodulation, [], [f4615, f2643])).
fof(f4615, plain, (~ (h18(op1(e13, e12)) = op2(h8(e11), e24)) | ~ (h8(e10) = h18(op1(e13, e13))) | ~ (h18(op1(e13, e14)) = op2(h8(e11), e21)) | ~ (h18(op1(e14, e10)) = op2(e21, h18(e10))) | ~ (h18(op1(e14, e11)) = op2(e21, h8(e13))) | ~ (h8(e13) = h18(op1(e14, e12))) | ~ (h18(op1(e14, e13)) = op2(e21, h8(e11))) | ~ (op2(e21, e21) = h18(op1(e14, e14))) | sP71 | sP70 | sP68 | ~ (h18(op1(e13, e11)) = op2(h18(e13), h18(e11))) | ~ (h18(op1(e13, e10)) = op2(h18(e13), h18(e10))) | ~ (h18(op1(e12, e14)) = op2(h18(e12), h18(e14))) | ~ (h18(op1(e12, e13)) = op2(h18(e12), h18(e13))) | ~ (h18(op1(e12, e12)) = op2(h18(e12), h18(e12))) | ~ (h18(op1(e12, e11)) = op2(h18(e12), h18(e11))) | ~ (h18(op1(e12, e10)) = op2(h18(e12), h18(e10))) | ~ (h18(op1(e11, e14)) = op2(h18(e11), h18(e14))) | ~ (h18(op1(e11, e13)) = op2(h18(e11), h18(e13))) | ~ (h18(op1(e11, e12)) = op2(h18(e11), h18(e12))) | ~ (h18(op1(e11, e11)) = op2(h18(e11), h18(e11))) | ~ (h18(op1(e11, e10)) = op2(h18(e11), h18(e10))) | ~ (h18(op1(e10, e14)) = op2(h18(e10), h18(e14))) | ~ (h18(op1(e10, e13)) = op2(h18(e10), h18(e13))) | ~ (h18(op1(e10, e12)) = op2(h18(e10), h18(e12))) | ~ (h18(op1(e10, e11)) = op2(h18(e10), h18(e11))) | ~ (h18(op1(e10, e10)) = op2(h18(e10), h18(e10)))), inference(forward_demodulation, [], [f4614, f2642])).
fof(f4614, plain, (~ (h18(op1(e13, e12)) = op2(h18(e13), e24)) | ~ (h8(e10) = h18(op1(e13, e13))) | ~ (h18(op1(e13, e14)) = op2(h8(e11), e21)) | ~ (h18(op1(e14, e10)) = op2(e21, h18(e10))) | ~ (h18(op1(e14, e11)) = op2(e21, h8(e13))) | ~ (h8(e13) = h18(op1(e14, e12))) | ~ (h18(op1(e14, e13)) = op2(e21, h8(e11))) | ~ (op2(e21, e21) = h18(op1(e14, e14))) | sP71 | sP70 | sP68 | ~ (h18(op1(e13, e11)) = op2(h18(e13), h18(e11))) | ~ (h18(op1(e13, e10)) = op2(h18(e13), h18(e10))) | ~ (h18(op1(e12, e14)) = op2(h18(e12), h18(e14))) | ~ (h18(op1(e12, e13)) = op2(h18(e12), h18(e13))) | ~ (h18(op1(e12, e12)) = op2(h18(e12), h18(e12))) | ~ (h18(op1(e12, e11)) = op2(h18(e12), h18(e11))) | ~ (h18(op1(e12, e10)) = op2(h18(e12), h18(e10))) | ~ (h18(op1(e11, e14)) = op2(h18(e11), h18(e14))) | ~ (h18(op1(e11, e13)) = op2(h18(e11), h18(e13))) | ~ (h18(op1(e11, e12)) = op2(h18(e11), h18(e12))) | ~ (h18(op1(e11, e11)) = op2(h18(e11), h18(e11))) | ~ (h18(op1(e11, e10)) = op2(h18(e11), h18(e10))) | ~ (h18(op1(e10, e14)) = op2(h18(e10), h18(e14))) | ~ (h18(op1(e10, e13)) = op2(h18(e10), h18(e13))) | ~ (h18(op1(e10, e12)) = op2(h18(e10), h18(e12))) | ~ (h18(op1(e10, e11)) = op2(h18(e10), h18(e11))) | ~ (h18(op1(e10, e10)) = op2(h18(e10), h18(e10)))), inference(forward_demodulation, [], [f4613, f758])).
fof(f4613, plain, (~ (h8(e10) = h18(op1(e13, e13))) | ~ (h18(op1(e13, e14)) = op2(h8(e11), e21)) | ~ (h18(op1(e14, e10)) = op2(e21, h18(e10))) | ~ (h18(op1(e14, e11)) = op2(e21, h8(e13))) | ~ (h8(e13) = h18(op1(e14, e12))) | ~ (h18(op1(e14, e13)) = op2(e21, h8(e11))) | ~ (op2(e21, e21) = h18(op1(e14, e14))) | sP71 | sP70 | sP68 | ~ (h18(op1(e13, e12)) = op2(h18(e13), h18(e12))) | ~ (h18(op1(e13, e11)) = op2(h18(e13), h18(e11))) | ~ (h18(op1(e13, e10)) = op2(h18(e13), h18(e10))) | ~ (h18(op1(e12, e14)) = op2(h18(e12), h18(e14))) | ~ (h18(op1(e12, e13)) = op2(h18(e12), h18(e13))) | ~ (h18(op1(e12, e12)) = op2(h18(e12), h18(e12))) | ~ (h18(op1(e12, e11)) = op2(h18(e12), h18(e11))) | ~ (h18(op1(e12, e10)) = op2(h18(e12), h18(e10))) | ~ (h18(op1(e11, e14)) = op2(h18(e11), h18(e14))) | ~ (h18(op1(e11, e13)) = op2(h18(e11), h18(e13))) | ~ (h18(op1(e11, e12)) = op2(h18(e11), h18(e12))) | ~ (h18(op1(e11, e11)) = op2(h18(e11), h18(e11))) | ~ (h18(op1(e11, e10)) = op2(h18(e11), h18(e10))) | ~ (h18(op1(e10, e14)) = op2(h18(e10), h18(e14))) | ~ (h18(op1(e10, e13)) = op2(h18(e10), h18(e13))) | ~ (h18(op1(e10, e12)) = op2(h18(e10), h18(e12))) | ~ (h18(op1(e10, e11)) = op2(h18(e10), h18(e11))) | ~ (h18(op1(e10, e10)) = op2(h18(e10), h18(e10)))), inference(forward_demodulation, [], [f4612, f2600])).
fof(f4612, plain, (~ (op2(h8(e11), h8(e11)) = h18(op1(e13, e13))) | ~ (h18(op1(e13, e14)) = op2(h8(e11), e21)) | ~ (h18(op1(e14, e10)) = op2(e21, h18(e10))) | ~ (h18(op1(e14, e11)) = op2(e21, h8(e13))) | ~ (h8(e13) = h18(op1(e14, e12))) | ~ (h18(op1(e14, e13)) = op2(e21, h8(e11))) | ~ (op2(e21, e21) = h18(op1(e14, e14))) | sP71 | sP70 | sP68 | ~ (h18(op1(e13, e12)) = op2(h18(e13), h18(e12))) | ~ (h18(op1(e13, e11)) = op2(h18(e13), h18(e11))) | ~ (h18(op1(e13, e10)) = op2(h18(e13), h18(e10))) | ~ (h18(op1(e12, e14)) = op2(h18(e12), h18(e14))) | ~ (h18(op1(e12, e13)) = op2(h18(e12), h18(e13))) | ~ (h18(op1(e12, e12)) = op2(h18(e12), h18(e12))) | ~ (h18(op1(e12, e11)) = op2(h18(e12), h18(e11))) | ~ (h18(op1(e12, e10)) = op2(h18(e12), h18(e10))) | ~ (h18(op1(e11, e14)) = op2(h18(e11), h18(e14))) | ~ (h18(op1(e11, e13)) = op2(h18(e11), h18(e13))) | ~ (h18(op1(e11, e12)) = op2(h18(e11), h18(e12))) | ~ (h18(op1(e11, e11)) = op2(h18(e11), h18(e11))) | ~ (h18(op1(e11, e10)) = op2(h18(e11), h18(e10))) | ~ (h18(op1(e10, e14)) = op2(h18(e10), h18(e14))) | ~ (h18(op1(e10, e13)) = op2(h18(e10), h18(e13))) | ~ (h18(op1(e10, e12)) = op2(h18(e10), h18(e12))) | ~ (h18(op1(e10, e11)) = op2(h18(e10), h18(e11))) | ~ (h18(op1(e10, e10)) = op2(h18(e10), h18(e10)))), inference(forward_demodulation, [], [f4611, f2642])).
fof(f4611, plain, (~ (h18(op1(e13, e14)) = op2(h8(e11), e21)) | ~ (h18(op1(e14, e10)) = op2(e21, h18(e10))) | ~ (h18(op1(e14, e11)) = op2(e21, h8(e13))) | ~ (h8(e13) = h18(op1(e14, e12))) | ~ (h18(op1(e14, e13)) = op2(e21, h8(e11))) | ~ (op2(e21, e21) = h18(op1(e14, e14))) | sP71 | sP70 | sP68 | ~ (h18(op1(e13, e13)) = op2(h18(e13), h18(e13))) | ~ (h18(op1(e13, e12)) = op2(h18(e13), h18(e12))) | ~ (h18(op1(e13, e11)) = op2(h18(e13), h18(e11))) | ~ (h18(op1(e13, e10)) = op2(h18(e13), h18(e10))) | ~ (h18(op1(e12, e14)) = op2(h18(e12), h18(e14))) | ~ (h18(op1(e12, e13)) = op2(h18(e12), h18(e13))) | ~ (h18(op1(e12, e12)) = op2(h18(e12), h18(e12))) | ~ (h18(op1(e12, e11)) = op2(h18(e12), h18(e11))) | ~ (h18(op1(e12, e10)) = op2(h18(e12), h18(e10))) | ~ (h18(op1(e11, e14)) = op2(h18(e11), h18(e14))) | ~ (h18(op1(e11, e13)) = op2(h18(e11), h18(e13))) | ~ (h18(op1(e11, e12)) = op2(h18(e11), h18(e12))) | ~ (h18(op1(e11, e11)) = op2(h18(e11), h18(e11))) | ~ (h18(op1(e11, e10)) = op2(h18(e11), h18(e10))) | ~ (h18(op1(e10, e14)) = op2(h18(e10), h18(e14))) | ~ (h18(op1(e10, e13)) = op2(h18(e10), h18(e13))) | ~ (h18(op1(e10, e12)) = op2(h18(e10), h18(e12))) | ~ (h18(op1(e10, e11)) = op2(h18(e10), h18(e11))) | ~ (h18(op1(e10, e10)) = op2(h18(e10), h18(e10)))), inference(forward_demodulation, [], [f4610, f2642])).
fof(f4610, plain, (~ (h18(op1(e13, e14)) = op2(h18(e13), e21)) | ~ (h18(op1(e14, e10)) = op2(e21, h18(e10))) | ~ (h18(op1(e14, e11)) = op2(e21, h8(e13))) | ~ (h8(e13) = h18(op1(e14, e12))) | ~ (h18(op1(e14, e13)) = op2(e21, h8(e11))) | ~ (op2(e21, e21) = h18(op1(e14, e14))) | sP71 | sP70 | sP68 | ~ (h18(op1(e13, e13)) = op2(h18(e13), h18(e13))) | ~ (h18(op1(e13, e12)) = op2(h18(e13), h18(e12))) | ~ (h18(op1(e13, e11)) = op2(h18(e13), h18(e11))) | ~ (h18(op1(e13, e10)) = op2(h18(e13), h18(e10))) | ~ (h18(op1(e12, e14)) = op2(h18(e12), h18(e14))) | ~ (h18(op1(e12, e13)) = op2(h18(e12), h18(e13))) | ~ (h18(op1(e12, e12)) = op2(h18(e12), h18(e12))) | ~ (h18(op1(e12, e11)) = op2(h18(e12), h18(e11))) | ~ (h18(op1(e12, e10)) = op2(h18(e12), h18(e10))) | ~ (h18(op1(e11, e14)) = op2(h18(e11), h18(e14))) | ~ (h18(op1(e11, e13)) = op2(h18(e11), h18(e13))) | ~ (h18(op1(e11, e12)) = op2(h18(e11), h18(e12))) | ~ (h18(op1(e11, e11)) = op2(h18(e11), h18(e11))) | ~ (h18(op1(e11, e10)) = op2(h18(e11), h18(e10))) | ~ (h18(op1(e10, e14)) = op2(h18(e10), h18(e14))) | ~ (h18(op1(e10, e13)) = op2(h18(e10), h18(e13))) | ~ (h18(op1(e10, e12)) = op2(h18(e10), h18(e12))) | ~ (h18(op1(e10, e11)) = op2(h18(e10), h18(e11))) | ~ (h18(op1(e10, e10)) = op2(h18(e10), h18(e10)))), inference(forward_demodulation, [], [f4609, f759])).
fof(f4609, plain, (~ (h18(op1(e14, e10)) = op2(e21, h18(e10))) | ~ (h18(op1(e14, e11)) = op2(e21, h8(e13))) | ~ (h8(e13) = h18(op1(e14, e12))) | ~ (h18(op1(e14, e13)) = op2(e21, h8(e11))) | ~ (op2(e21, e21) = h18(op1(e14, e14))) | sP71 | sP70 | sP68 | ~ (h18(op1(e13, e14)) = op2(h18(e13), h18(e14))) | ~ (h18(op1(e13, e13)) = op2(h18(e13), h18(e13))) | ~ (h18(op1(e13, e12)) = op2(h18(e13), h18(e12))) | ~ (h18(op1(e13, e11)) = op2(h18(e13), h18(e11))) | ~ (h18(op1(e13, e10)) = op2(h18(e13), h18(e10))) | ~ (h18(op1(e12, e14)) = op2(h18(e12), h18(e14))) | ~ (h18(op1(e12, e13)) = op2(h18(e12), h18(e13))) | ~ (h18(op1(e12, e12)) = op2(h18(e12), h18(e12))) | ~ (h18(op1(e12, e11)) = op2(h18(e12), h18(e11))) | ~ (h18(op1(e12, e10)) = op2(h18(e12), h18(e10))) | ~ (h18(op1(e11, e14)) = op2(h18(e11), h18(e14))) | ~ (h18(op1(e11, e13)) = op2(h18(e11), h18(e13))) | ~ (h18(op1(e11, e12)) = op2(h18(e11), h18(e12))) | ~ (h18(op1(e11, e11)) = op2(h18(e11), h18(e11))) | ~ (h18(op1(e11, e10)) = op2(h18(e11), h18(e10))) | ~ (h18(op1(e10, e14)) = op2(h18(e10), h18(e14))) | ~ (h18(op1(e10, e13)) = op2(h18(e10), h18(e13))) | ~ (h18(op1(e10, e12)) = op2(h18(e10), h18(e12))) | ~ (h18(op1(e10, e11)) = op2(h18(e10), h18(e11))) | ~ (h18(op1(e10, e10)) = op2(h18(e10), h18(e10)))), inference(forward_demodulation, [], [f4608, f759])).
fof(f4608, plain, (~ (h18(op1(e14, e11)) = op2(e21, h8(e13))) | ~ (h8(e13) = h18(op1(e14, e12))) | ~ (h18(op1(e14, e13)) = op2(e21, h8(e11))) | ~ (op2(e21, e21) = h18(op1(e14, e14))) | sP71 | sP70 | sP68 | ~ (h18(op1(e14, e10)) = op2(h18(e14), h18(e10))) | ~ (h18(op1(e13, e14)) = op2(h18(e13), h18(e14))) | ~ (h18(op1(e13, e13)) = op2(h18(e13), h18(e13))) | ~ (h18(op1(e13, e12)) = op2(h18(e13), h18(e12))) | ~ (h18(op1(e13, e11)) = op2(h18(e13), h18(e11))) | ~ (h18(op1(e13, e10)) = op2(h18(e13), h18(e10))) | ~ (h18(op1(e12, e14)) = op2(h18(e12), h18(e14))) | ~ (h18(op1(e12, e13)) = op2(h18(e12), h18(e13))) | ~ (h18(op1(e12, e12)) = op2(h18(e12), h18(e12))) | ~ (h18(op1(e12, e11)) = op2(h18(e12), h18(e11))) | ~ (h18(op1(e12, e10)) = op2(h18(e12), h18(e10))) | ~ (h18(op1(e11, e14)) = op2(h18(e11), h18(e14))) | ~ (h18(op1(e11, e13)) = op2(h18(e11), h18(e13))) | ~ (h18(op1(e11, e12)) = op2(h18(e11), h18(e12))) | ~ (h18(op1(e11, e11)) = op2(h18(e11), h18(e11))) | ~ (h18(op1(e11, e10)) = op2(h18(e11), h18(e10))) | ~ (h18(op1(e10, e14)) = op2(h18(e10), h18(e14))) | ~ (h18(op1(e10, e13)) = op2(h18(e10), h18(e13))) | ~ (h18(op1(e10, e12)) = op2(h18(e10), h18(e12))) | ~ (h18(op1(e10, e11)) = op2(h18(e10), h18(e11))) | ~ (h18(op1(e10, e10)) = op2(h18(e10), h18(e10)))), inference(forward_demodulation, [], [f4607, f759])).
fof(f4607, plain, (~ (h18(op1(e14, e11)) = op2(h18(e14), h8(e13))) | ~ (h8(e13) = h18(op1(e14, e12))) | ~ (h18(op1(e14, e13)) = op2(e21, h8(e11))) | ~ (op2(e21, e21) = h18(op1(e14, e14))) | sP71 | sP70 | sP68 | ~ (h18(op1(e14, e10)) = op2(h18(e14), h18(e10))) | ~ (h18(op1(e13, e14)) = op2(h18(e13), h18(e14))) | ~ (h18(op1(e13, e13)) = op2(h18(e13), h18(e13))) | ~ (h18(op1(e13, e12)) = op2(h18(e13), h18(e12))) | ~ (h18(op1(e13, e11)) = op2(h18(e13), h18(e11))) | ~ (h18(op1(e13, e10)) = op2(h18(e13), h18(e10))) | ~ (h18(op1(e12, e14)) = op2(h18(e12), h18(e14))) | ~ (h18(op1(e12, e13)) = op2(h18(e12), h18(e13))) | ~ (h18(op1(e12, e12)) = op2(h18(e12), h18(e12))) | ~ (h18(op1(e12, e11)) = op2(h18(e12), h18(e11))) | ~ (h18(op1(e12, e10)) = op2(h18(e12), h18(e10))) | ~ (h18(op1(e11, e14)) = op2(h18(e11), h18(e14))) | ~ (h18(op1(e11, e13)) = op2(h18(e11), h18(e13))) | ~ (h18(op1(e11, e12)) = op2(h18(e11), h18(e12))) | ~ (h18(op1(e11, e11)) = op2(h18(e11), h18(e11))) | ~ (h18(op1(e11, e10)) = op2(h18(e11), h18(e10))) | ~ (h18(op1(e10, e14)) = op2(h18(e10), h18(e14))) | ~ (h18(op1(e10, e13)) = op2(h18(e10), h18(e13))) | ~ (h18(op1(e10, e12)) = op2(h18(e10), h18(e12))) | ~ (h18(op1(e10, e11)) = op2(h18(e10), h18(e11))) | ~ (h18(op1(e10, e10)) = op2(h18(e10), h18(e10)))), inference(forward_demodulation, [], [f4606, f2643])).
fof(f4606, plain, (~ (h8(e13) = h18(op1(e14, e12))) | ~ (h18(op1(e14, e13)) = op2(e21, h8(e11))) | ~ (op2(e21, e21) = h18(op1(e14, e14))) | sP71 | sP70 | sP68 | ~ (h18(op1(e14, e11)) = op2(h18(e14), h18(e11))) | ~ (h18(op1(e14, e10)) = op2(h18(e14), h18(e10))) | ~ (h18(op1(e13, e14)) = op2(h18(e13), h18(e14))) | ~ (h18(op1(e13, e13)) = op2(h18(e13), h18(e13))) | ~ (h18(op1(e13, e12)) = op2(h18(e13), h18(e12))) | ~ (h18(op1(e13, e11)) = op2(h18(e13), h18(e11))) | ~ (h18(op1(e13, e10)) = op2(h18(e13), h18(e10))) | ~ (h18(op1(e12, e14)) = op2(h18(e12), h18(e14))) | ~ (h18(op1(e12, e13)) = op2(h18(e12), h18(e13))) | ~ (h18(op1(e12, e12)) = op2(h18(e12), h18(e12))) | ~ (h18(op1(e12, e11)) = op2(h18(e12), h18(e11))) | ~ (h18(op1(e12, e10)) = op2(h18(e12), h18(e10))) | ~ (h18(op1(e11, e14)) = op2(h18(e11), h18(e14))) | ~ (h18(op1(e11, e13)) = op2(h18(e11), h18(e13))) | ~ (h18(op1(e11, e12)) = op2(h18(e11), h18(e12))) | ~ (h18(op1(e11, e11)) = op2(h18(e11), h18(e11))) | ~ (h18(op1(e11, e10)) = op2(h18(e11), h18(e10))) | ~ (h18(op1(e10, e14)) = op2(h18(e10), h18(e14))) | ~ (h18(op1(e10, e13)) = op2(h18(e10), h18(e13))) | ~ (h18(op1(e10, e12)) = op2(h18(e10), h18(e12))) | ~ (h18(op1(e10, e11)) = op2(h18(e10), h18(e11))) | ~ (h18(op1(e10, e10)) = op2(h18(e10), h18(e10)))), inference(forward_demodulation, [], [f4605, f712])).
fof(f4605, plain, (~ (op2(e21, e24) = h18(op1(e14, e12))) | ~ (h18(op1(e14, e13)) = op2(e21, h8(e11))) | ~ (op2(e21, e21) = h18(op1(e14, e14))) | sP71 | sP70 | sP68 | ~ (h18(op1(e14, e11)) = op2(h18(e14), h18(e11))) | ~ (h18(op1(e14, e10)) = op2(h18(e14), h18(e10))) | ~ (h18(op1(e13, e14)) = op2(h18(e13), h18(e14))) | ~ (h18(op1(e13, e13)) = op2(h18(e13), h18(e13))) | ~ (h18(op1(e13, e12)) = op2(h18(e13), h18(e12))) | ~ (h18(op1(e13, e11)) = op2(h18(e13), h18(e11))) | ~ (h18(op1(e13, e10)) = op2(h18(e13), h18(e10))) | ~ (h18(op1(e12, e14)) = op2(h18(e12), h18(e14))) | ~ (h18(op1(e12, e13)) = op2(h18(e12), h18(e13))) | ~ (h18(op1(e12, e12)) = op2(h18(e12), h18(e12))) | ~ (h18(op1(e12, e11)) = op2(h18(e12), h18(e11))) | ~ (h18(op1(e12, e10)) = op2(h18(e12), h18(e10))) | ~ (h18(op1(e11, e14)) = op2(h18(e11), h18(e14))) | ~ (h18(op1(e11, e13)) = op2(h18(e11), h18(e13))) | ~ (h18(op1(e11, e12)) = op2(h18(e11), h18(e12))) | ~ (h18(op1(e11, e11)) = op2(h18(e11), h18(e11))) | ~ (h18(op1(e11, e10)) = op2(h18(e11), h18(e10))) | ~ (h18(op1(e10, e14)) = op2(h18(e10), h18(e14))) | ~ (h18(op1(e10, e13)) = op2(h18(e10), h18(e13))) | ~ (h18(op1(e10, e12)) = op2(h18(e10), h18(e12))) | ~ (h18(op1(e10, e11)) = op2(h18(e10), h18(e11))) | ~ (h18(op1(e10, e10)) = op2(h18(e10), h18(e10)))), inference(forward_demodulation, [], [f4604, f759])).
fof(f4604, plain, (~ (h18(op1(e14, e12)) = op2(h18(e14), e24)) | ~ (h18(op1(e14, e13)) = op2(e21, h8(e11))) | ~ (op2(e21, e21) = h18(op1(e14, e14))) | sP71 | sP70 | sP68 | ~ (h18(op1(e14, e11)) = op2(h18(e14), h18(e11))) | ~ (h18(op1(e14, e10)) = op2(h18(e14), h18(e10))) | ~ (h18(op1(e13, e14)) = op2(h18(e13), h18(e14))) | ~ (h18(op1(e13, e13)) = op2(h18(e13), h18(e13))) | ~ (h18(op1(e13, e12)) = op2(h18(e13), h18(e12))) | ~ (h18(op1(e13, e11)) = op2(h18(e13), h18(e11))) | ~ (h18(op1(e13, e10)) = op2(h18(e13), h18(e10))) | ~ (h18(op1(e12, e14)) = op2(h18(e12), h18(e14))) | ~ (h18(op1(e12, e13)) = op2(h18(e12), h18(e13))) | ~ (h18(op1(e12, e12)) = op2(h18(e12), h18(e12))) | ~ (h18(op1(e12, e11)) = op2(h18(e12), h18(e11))) | ~ (h18(op1(e12, e10)) = op2(h18(e12), h18(e10))) | ~ (h18(op1(e11, e14)) = op2(h18(e11), h18(e14))) | ~ (h18(op1(e11, e13)) = op2(h18(e11), h18(e13))) | ~ (h18(op1(e11, e12)) = op2(h18(e11), h18(e12))) | ~ (h18(op1(e11, e11)) = op2(h18(e11), h18(e11))) | ~ (h18(op1(e11, e10)) = op2(h18(e11), h18(e10))) | ~ (h18(op1(e10, e14)) = op2(h18(e10), h18(e14))) | ~ (h18(op1(e10, e13)) = op2(h18(e10), h18(e13))) | ~ (h18(op1(e10, e12)) = op2(h18(e10), h18(e12))) | ~ (h18(op1(e10, e11)) = op2(h18(e10), h18(e11))) | ~ (h18(op1(e10, e10)) = op2(h18(e10), h18(e10)))), inference(forward_demodulation, [], [f4603, f758])).
fof(f4603, plain, (~ (h18(op1(e14, e13)) = op2(e21, h8(e11))) | ~ (op2(e21, e21) = h18(op1(e14, e14))) | sP71 | sP70 | sP68 | ~ (h18(op1(e14, e12)) = op2(h18(e14), h18(e12))) | ~ (h18(op1(e14, e11)) = op2(h18(e14), h18(e11))) | ~ (h18(op1(e14, e10)) = op2(h18(e14), h18(e10))) | ~ (h18(op1(e13, e14)) = op2(h18(e13), h18(e14))) | ~ (h18(op1(e13, e13)) = op2(h18(e13), h18(e13))) | ~ (h18(op1(e13, e12)) = op2(h18(e13), h18(e12))) | ~ (h18(op1(e13, e11)) = op2(h18(e13), h18(e11))) | ~ (h18(op1(e13, e10)) = op2(h18(e13), h18(e10))) | ~ (h18(op1(e12, e14)) = op2(h18(e12), h18(e14))) | ~ (h18(op1(e12, e13)) = op2(h18(e12), h18(e13))) | ~ (h18(op1(e12, e12)) = op2(h18(e12), h18(e12))) | ~ (h18(op1(e12, e11)) = op2(h18(e12), h18(e11))) | ~ (h18(op1(e12, e10)) = op2(h18(e12), h18(e10))) | ~ (h18(op1(e11, e14)) = op2(h18(e11), h18(e14))) | ~ (h18(op1(e11, e13)) = op2(h18(e11), h18(e13))) | ~ (h18(op1(e11, e12)) = op2(h18(e11), h18(e12))) | ~ (h18(op1(e11, e11)) = op2(h18(e11), h18(e11))) | ~ (h18(op1(e11, e10)) = op2(h18(e11), h18(e10))) | ~ (h18(op1(e10, e14)) = op2(h18(e10), h18(e14))) | ~ (h18(op1(e10, e13)) = op2(h18(e10), h18(e13))) | ~ (h18(op1(e10, e12)) = op2(h18(e10), h18(e12))) | ~ (h18(op1(e10, e11)) = op2(h18(e10), h18(e11))) | ~ (h18(op1(e10, e10)) = op2(h18(e10), h18(e10)))), inference(forward_demodulation, [], [f4602, f759])).
fof(f4602, plain, (~ (h18(op1(e14, e13)) = op2(h18(e14), h8(e11))) | ~ (op2(e21, e21) = h18(op1(e14, e14))) | sP71 | sP70 | sP68 | ~ (h18(op1(e14, e12)) = op2(h18(e14), h18(e12))) | ~ (h18(op1(e14, e11)) = op2(h18(e14), h18(e11))) | ~ (h18(op1(e14, e10)) = op2(h18(e14), h18(e10))) | ~ (h18(op1(e13, e14)) = op2(h18(e13), h18(e14))) | ~ (h18(op1(e13, e13)) = op2(h18(e13), h18(e13))) | ~ (h18(op1(e13, e12)) = op2(h18(e13), h18(e12))) | ~ (h18(op1(e13, e11)) = op2(h18(e13), h18(e11))) | ~ (h18(op1(e13, e10)) = op2(h18(e13), h18(e10))) | ~ (h18(op1(e12, e14)) = op2(h18(e12), h18(e14))) | ~ (h18(op1(e12, e13)) = op2(h18(e12), h18(e13))) | ~ (h18(op1(e12, e12)) = op2(h18(e12), h18(e12))) | ~ (h18(op1(e12, e11)) = op2(h18(e12), h18(e11))) | ~ (h18(op1(e12, e10)) = op2(h18(e12), h18(e10))) | ~ (h18(op1(e11, e14)) = op2(h18(e11), h18(e14))) | ~ (h18(op1(e11, e13)) = op2(h18(e11), h18(e13))) | ~ (h18(op1(e11, e12)) = op2(h18(e11), h18(e12))) | ~ (h18(op1(e11, e11)) = op2(h18(e11), h18(e11))) | ~ (h18(op1(e11, e10)) = op2(h18(e11), h18(e10))) | ~ (h18(op1(e10, e14)) = op2(h18(e10), h18(e14))) | ~ (h18(op1(e10, e13)) = op2(h18(e10), h18(e13))) | ~ (h18(op1(e10, e12)) = op2(h18(e10), h18(e12))) | ~ (h18(op1(e10, e11)) = op2(h18(e10), h18(e11))) | ~ (h18(op1(e10, e10)) = op2(h18(e10), h18(e10)))), inference(forward_demodulation, [], [f4601, f2642])).
fof(f4601, plain, (~ (op2(e21, e21) = h18(op1(e14, e14))) | sP71 | sP70 | sP68 | ~ (h18(op1(e14, e13)) = op2(h18(e14), h18(e13))) | ~ (h18(op1(e14, e12)) = op2(h18(e14), h18(e12))) | ~ (h18(op1(e14, e11)) = op2(h18(e14), h18(e11))) | ~ (h18(op1(e14, e10)) = op2(h18(e14), h18(e10))) | ~ (h18(op1(e13, e14)) = op2(h18(e13), h18(e14))) | ~ (h18(op1(e13, e13)) = op2(h18(e13), h18(e13))) | ~ (h18(op1(e13, e12)) = op2(h18(e13), h18(e12))) | ~ (h18(op1(e13, e11)) = op2(h18(e13), h18(e11))) | ~ (h18(op1(e13, e10)) = op2(h18(e13), h18(e10))) | ~ (h18(op1(e12, e14)) = op2(h18(e12), h18(e14))) | ~ (h18(op1(e12, e13)) = op2(h18(e12), h18(e13))) | ~ (h18(op1(e12, e12)) = op2(h18(e12), h18(e12))) | ~ (h18(op1(e12, e11)) = op2(h18(e12), h18(e11))) | ~ (h18(op1(e12, e10)) = op2(h18(e12), h18(e10))) | ~ (h18(op1(e11, e14)) = op2(h18(e11), h18(e14))) | ~ (h18(op1(e11, e13)) = op2(h18(e11), h18(e13))) | ~ (h18(op1(e11, e12)) = op2(h18(e11), h18(e12))) | ~ (h18(op1(e11, e11)) = op2(h18(e11), h18(e11))) | ~ (h18(op1(e11, e10)) = op2(h18(e11), h18(e10))) | ~ (h18(op1(e10, e14)) = op2(h18(e10), h18(e14))) | ~ (h18(op1(e10, e13)) = op2(h18(e10), h18(e13))) | ~ (h18(op1(e10, e12)) = op2(h18(e10), h18(e12))) | ~ (h18(op1(e10, e11)) = op2(h18(e10), h18(e11))) | ~ (h18(op1(e10, e10)) = op2(h18(e10), h18(e10)))), inference(forward_demodulation, [], [f4600, f759])).
fof(f4600, plain, (sP71 | sP70 | sP68 | ~ (h18(op1(e14, e14)) = op2(h18(e14), h18(e14))) | ~ (h18(op1(e14, e13)) = op2(h18(e14), h18(e13))) | ~ (h18(op1(e14, e12)) = op2(h18(e14), h18(e12))) | ~ (h18(op1(e14, e11)) = op2(h18(e14), h18(e11))) | ~ (h18(op1(e14, e10)) = op2(h18(e14), h18(e10))) | ~ (h18(op1(e13, e14)) = op2(h18(e13), h18(e14))) | ~ (h18(op1(e13, e13)) = op2(h18(e13), h18(e13))) | ~ (h18(op1(e13, e12)) = op2(h18(e13), h18(e12))) | ~ (h18(op1(e13, e11)) = op2(h18(e13), h18(e11))) | ~ (h18(op1(e13, e10)) = op2(h18(e13), h18(e10))) | ~ (h18(op1(e12, e14)) = op2(h18(e12), h18(e14))) | ~ (h18(op1(e12, e13)) = op2(h18(e12), h18(e13))) | ~ (h18(op1(e12, e12)) = op2(h18(e12), h18(e12))) | ~ (h18(op1(e12, e11)) = op2(h18(e12), h18(e11))) | ~ (h18(op1(e12, e10)) = op2(h18(e12), h18(e10))) | ~ (h18(op1(e11, e14)) = op2(h18(e11), h18(e14))) | ~ (h18(op1(e11, e13)) = op2(h18(e11), h18(e13))) | ~ (h18(op1(e11, e12)) = op2(h18(e11), h18(e12))) | ~ (h18(op1(e11, e11)) = op2(h18(e11), h18(e11))) | ~ (h18(op1(e11, e10)) = op2(h18(e11), h18(e10))) | ~ (h18(op1(e10, e14)) = op2(h18(e10), h18(e14))) | ~ (h18(op1(e10, e13)) = op2(h18(e10), h18(e13))) | ~ (h18(op1(e10, e12)) = op2(h18(e10), h18(e12))) | ~ (h18(op1(e10, e11)) = op2(h18(e10), h18(e11))) | ~ (h18(op1(e10, e10)) = op2(h18(e10), h18(e10)))), inference(subsumption_resolution, [], [f4599, f2837])).
fof(f2837, plain, ~ sP69, inference(subsumption_resolution, [], [f827, f759])).
fof(f827, plain, (~ (e21 = h18(e14)) | ~ sP69), inference(cnf_transformation, [], [f130])).
fof(f130, plain, ((~ (e21 = h18(e14)) & ~ (e21 = h18(e13)) & ~ (e21 = h18(e12)) & ~ (e21 = h18(e11)) & ~ (e21 = h18(e10))) | ~ sP69), inference(nnf_transformation, [], [f108])).
fof(f108, plain, ((~ (e21 = h18(e14)) & ~ (e21 = h18(e13)) & ~ (e21 = h18(e12)) & ~ (e21 = h18(e11)) & ~ (e21 = h18(e10))) | ~ sP69), inference(usedef, [], [e108])).
fof(e108, plain, (sP69 <=> (~ (e21 = h18(e14)) & ~ (e21 = h18(e13)) & ~ (e21 = h18(e12)) & ~ (e21 = h18(e11)) & ~ (e21 = h18(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP69])])).
fof(f4599, plain, (sP71 | sP70 | sP69 | sP68 | ~ (h18(op1(e14, e14)) = op2(h18(e14), h18(e14))) | ~ (h18(op1(e14, e13)) = op2(h18(e14), h18(e13))) | ~ (h18(op1(e14, e12)) = op2(h18(e14), h18(e12))) | ~ (h18(op1(e14, e11)) = op2(h18(e14), h18(e11))) | ~ (h18(op1(e14, e10)) = op2(h18(e14), h18(e10))) | ~ (h18(op1(e13, e14)) = op2(h18(e13), h18(e14))) | ~ (h18(op1(e13, e13)) = op2(h18(e13), h18(e13))) | ~ (h18(op1(e13, e12)) = op2(h18(e13), h18(e12))) | ~ (h18(op1(e13, e11)) = op2(h18(e13), h18(e11))) | ~ (h18(op1(e13, e10)) = op2(h18(e13), h18(e10))) | ~ (h18(op1(e12, e14)) = op2(h18(e12), h18(e14))) | ~ (h18(op1(e12, e13)) = op2(h18(e12), h18(e13))) | ~ (h18(op1(e12, e12)) = op2(h18(e12), h18(e12))) | ~ (h18(op1(e12, e11)) = op2(h18(e12), h18(e11))) | ~ (h18(op1(e12, e10)) = op2(h18(e12), h18(e10))) | ~ (h18(op1(e11, e14)) = op2(h18(e11), h18(e14))) | ~ (h18(op1(e11, e13)) = op2(h18(e11), h18(e13))) | ~ (h18(op1(e11, e12)) = op2(h18(e11), h18(e12))) | ~ (h18(op1(e11, e11)) = op2(h18(e11), h18(e11))) | ~ (h18(op1(e11, e10)) = op2(h18(e11), h18(e10))) | ~ (h18(op1(e10, e14)) = op2(h18(e10), h18(e14))) | ~ (h18(op1(e10, e13)) = op2(h18(e10), h18(e13))) | ~ (h18(op1(e10, e12)) = op2(h18(e10), h18(e12))) | ~ (h18(op1(e10, e11)) = op2(h18(e10), h18(e11))) | ~ (h18(op1(e10, e10)) = op2(h18(e10), h18(e10)))), inference(subsumption_resolution, [], [f1260, f758])).
fof(f1260, plain, (~ (e24 = h18(e12)) | sP71 | sP70 | sP69 | sP68 | ~ (h18(op1(e14, e14)) = op2(h18(e14), h18(e14))) | ~ (h18(op1(e14, e13)) = op2(h18(e14), h18(e13))) | ~ (h18(op1(e14, e12)) = op2(h18(e14), h18(e12))) | ~ (h18(op1(e14, e11)) = op2(h18(e14), h18(e11))) | ~ (h18(op1(e14, e10)) = op2(h18(e14), h18(e10))) | ~ (h18(op1(e13, e14)) = op2(h18(e13), h18(e14))) | ~ (h18(op1(e13, e13)) = op2(h18(e13), h18(e13))) | ~ (h18(op1(e13, e12)) = op2(h18(e13), h18(e12))) | ~ (h18(op1(e13, e11)) = op2(h18(e13), h18(e11))) | ~ (h18(op1(e13, e10)) = op2(h18(e13), h18(e10))) | ~ (h18(op1(e12, e14)) = op2(h18(e12), h18(e14))) | ~ (h18(op1(e12, e13)) = op2(h18(e12), h18(e13))) | ~ (h18(op1(e12, e12)) = op2(h18(e12), h18(e12))) | ~ (h18(op1(e12, e11)) = op2(h18(e12), h18(e11))) | ~ (h18(op1(e12, e10)) = op2(h18(e12), h18(e10))) | ~ (h18(op1(e11, e14)) = op2(h18(e11), h18(e14))) | ~ (h18(op1(e11, e13)) = op2(h18(e11), h18(e13))) | ~ (h18(op1(e11, e12)) = op2(h18(e11), h18(e12))) | ~ (h18(op1(e11, e11)) = op2(h18(e11), h18(e11))) | ~ (h18(op1(e11, e10)) = op2(h18(e11), h18(e10))) | ~ (h18(op1(e10, e14)) = op2(h18(e10), h18(e14))) | ~ (h18(op1(e10, e13)) = op2(h18(e10), h18(e13))) | ~ (h18(op1(e10, e12)) = op2(h18(e10), h18(e12))) | ~ (h18(op1(e10, e11)) = op2(h18(e10), h18(e11))) | ~ (h18(op1(e10, e10)) = op2(h18(e10), h18(e10)))), inference(cnf_transformation, [], [f119])).
fof(f119, plain, (((~ (e24 = h20(e14)) & ~ (e24 = h20(e13)) & ~ (e24 = h20(e12)) & ~ (e24 = h20(e11)) & ~ (e24 = h20(e10))) | sP79 | sP78 | sP77 | sP76 | ~ (h20(op1(e14, e14)) = op2(h20(e14), h20(e14))) | ~ (h20(op1(e14, e13)) = op2(h20(e14), h20(e13))) | ~ (h20(op1(e14, e12)) = op2(h20(e14), h20(e12))) | ~ (h20(op1(e14, e11)) = op2(h20(e14), h20(e11))) | ~ (h20(op1(e14, e10)) = op2(h20(e14), h20(e10))) | ~ (h20(op1(e13, e14)) = op2(h20(e13), h20(e14))) | ~ (h20(op1(e13, e13)) = op2(h20(e13), h20(e13))) | ~ (h20(op1(e13, e12)) = op2(h20(e13), h20(e12))) | ~ (h20(op1(e13, e11)) = op2(h20(e13), h20(e11))) | ~ (h20(op1(e13, e10)) = op2(h20(e13), h20(e10))) | ~ (h20(op1(e12, e14)) = op2(h20(e12), h20(e14))) | ~ (h20(op1(e12, e13)) = op2(h20(e12), h20(e13))) | ~ (h20(op1(e12, e12)) = op2(h20(e12), h20(e12))) | ~ (h20(op1(e12, e11)) = op2(h20(e12), h20(e11))) | ~ (h20(op1(e12, e10)) = op2(h20(e12), h20(e10))) | ~ (h20(op1(e11, e14)) = op2(h20(e11), h20(e14))) | ~ (h20(op1(e11, e13)) = op2(h20(e11), h20(e13))) | ~ (h20(op1(e11, e12)) = op2(h20(e11), h20(e12))) | ~ (h20(op1(e11, e11)) = op2(h20(e11), h20(e11))) | ~ (h20(op1(e11, e10)) = op2(h20(e11), h20(e10))) | ~ (h20(op1(e10, e14)) = op2(h20(e10), h20(e14))) | ~ (h20(op1(e10, e13)) = op2(h20(e10), h20(e13))) | ~ (h20(op1(e10, e12)) = op2(h20(e10), h20(e12))) | ~ (h20(op1(e10, e11)) = op2(h20(e10), h20(e11))) | ~ (h20(op1(e10, e10)) = op2(h20(e10), h20(e10)))) & ((~ (e24 = h19(e14)) & ~ (e24 = h19(e13)) & ~ (e24 = h19(e12)) & ~ (e24 = h19(e11)) & ~ (e24 = h19(e10))) | sP75 | sP74 | sP73 | sP72 | ~ (h19(op1(e14, e14)) = op2(h19(e14), h19(e14))) | ~ (h19(op1(e14, e13)) = op2(h19(e14), h19(e13))) | ~ (h19(op1(e14, e12)) = op2(h19(e14), h19(e12))) | ~ (h19(op1(e14, e11)) = op2(h19(e14), h19(e11))) | ~ (h19(op1(e14, e10)) = op2(h19(e14), h19(e10))) | ~ (h19(op1(e13, e14)) = op2(h19(e13), h19(e14))) | ~ (h19(op1(e13, e13)) = op2(h19(e13), h19(e13))) | ~ (h19(op1(e13, e12)) = op2(h19(e13), h19(e12))) | ~ (h19(op1(e13, e11)) = op2(h19(e13), h19(e11))) | ~ (h19(op1(e13, e10)) = op2(h19(e13), h19(e10))) | ~ (h19(op1(e12, e14)) = op2(h19(e12), h19(e14))) | ~ (h19(op1(e12, e13)) = op2(h19(e12), h19(e13))) | ~ (h19(op1(e12, e12)) = op2(h19(e12), h19(e12))) | ~ (h19(op1(e12, e11)) = op2(h19(e12), h19(e11))) | ~ (h19(op1(e12, e10)) = op2(h19(e12), h19(e10))) | ~ (h19(op1(e11, e14)) = op2(h19(e11), h19(e14))) | ~ (h19(op1(e11, e13)) = op2(h19(e11), h19(e13))) | ~ (h19(op1(e11, e12)) = op2(h19(e11), h19(e12))) | ~ (h19(op1(e11, e11)) = op2(h19(e11), h19(e11))) | ~ (h19(op1(e11, e10)) = op2(h19(e11), h19(e10))) | ~ (h19(op1(e10, e14)) = op2(h19(e10), h19(e14))) | ~ (h19(op1(e10, e13)) = op2(h19(e10), h19(e13))) | ~ (h19(op1(e10, e12)) = op2(h19(e10), h19(e12))) | ~ (h19(op1(e10, e11)) = op2(h19(e10), h19(e11))) | ~ (h19(op1(e10, e10)) = op2(h19(e10), h19(e10)))) & ((~ (e24 = h18(e14)) & ~ (e24 = h18(e13)) & ~ (e24 = h18(e12)) & ~ (e24 = h18(e11)) & ~ (e24 = h18(e10))) | sP71 | sP70 | sP69 | sP68 | ~ (h18(op1(e14, e14)) = op2(h18(e14), h18(e14))) | ~ (h18(op1(e14, e13)) = op2(h18(e14), h18(e13))) | ~ (h18(op1(e14, e12)) = op2(h18(e14), h18(e12))) | ~ (h18(op1(e14, e11)) = op2(h18(e14), h18(e11))) | ~ (h18(op1(e14, e10)) = op2(h18(e14), h18(e10))) | ~ (h18(op1(e13, e14)) = op2(h18(e13), h18(e14))) | ~ (h18(op1(e13, e13)) = op2(h18(e13), h18(e13))) | ~ (h18(op1(e13, e12)) = op2(h18(e13), h18(e12))) | ~ (h18(op1(e13, e11)) = op2(h18(e13), h18(e11))) | ~ (h18(op1(e13, e10)) = op2(h18(e13), h18(e10))) | ~ (h18(op1(e12, e14)) = op2(h18(e12), h18(e14))) | ~ (h18(op1(e12, e13)) = op2(h18(e12), h18(e13))) | ~ (h18(op1(e12, e12)) = op2(h18(e12), h18(e12))) | ~ (h18(op1(e12, e11)) = op2(h18(e12), h18(e11))) | ~ (h18(op1(e12, e10)) = op2(h18(e12), h18(e10))) | ~ (h18(op1(e11, e14)) = op2(h18(e11), h18(e14))) | ~ (h18(op1(e11, e13)) = op2(h18(e11), h18(e13))) | ~ (h18(op1(e11, e12)) = op2(h18(e11), h18(e12))) | ~ (h18(op1(e11, e11)) = op2(h18(e11), h18(e11))) | ~ (h18(op1(e11, e10)) = op2(h18(e11), h18(e10))) | ~ (h18(op1(e10, e14)) = op2(h18(e10), h18(e14))) | ~ (h18(op1(e10, e13)) = op2(h18(e10), h18(e13))) | ~ (h18(op1(e10, e12)) = op2(h18(e10), h18(e12))) | ~ (h18(op1(e10, e11)) = op2(h18(e10), h18(e11))) | ~ (h18(op1(e10, e10)) = op2(h18(e10), h18(e10)))) & ((~ (e24 = h17(e14)) & ~ (e24 = h17(e13)) & ~ (e24 = h17(e12)) & ~ (e24 = h17(e11)) & ~ (e24 = h17(e10))) | sP67 | sP66 | sP65 | sP64 | ~ (h17(op1(e14, e14)) = op2(h17(e14), h17(e14))) | ~ (h17(op1(e14, e13)) = op2(h17(e14), h17(e13))) | ~ (h17(op1(e14, e12)) = op2(h17(e14), h17(e12))) | ~ (h17(op1(e14, e11)) = op2(h17(e14), h17(e11))) | ~ (h17(op1(e14, e10)) = op2(h17(e14), h17(e10))) | ~ (h17(op1(e13, e14)) = op2(h17(e13), h17(e14))) | ~ (h17(op1(e13, e13)) = op2(h17(e13), h17(e13))) | ~ (h17(op1(e13, e12)) = op2(h17(e13), h17(e12))) | ~ (h17(op1(e13, e11)) = op2(h17(e13), h17(e11))) | ~ (h17(op1(e13, e10)) = op2(h17(e13), h17(e10))) | ~ (h17(op1(e12, e14)) = op2(h17(e12), h17(e14))) | ~ (h17(op1(e12, e13)) = op2(h17(e12), h17(e13))) | ~ (h17(op1(e12, e12)) = op2(h17(e12), h17(e12))) | ~ (h17(op1(e12, e11)) = op2(h17(e12), h17(e11))) | ~ (h17(op1(e12, e10)) = op2(h17(e12), h17(e10))) | ~ (h17(op1(e11, e14)) = op2(h17(e11), h17(e14))) | ~ (h17(op1(e11, e13)) = op2(h17(e11), h17(e13))) | ~ (h17(op1(e11, e12)) = op2(h17(e11), h17(e12))) | ~ (h17(op1(e11, e11)) = op2(h17(e11), h17(e11))) | ~ (h17(op1(e11, e10)) = op2(h17(e11), h17(e10))) | ~ (h17(op1(e10, e14)) = op2(h17(e10), h17(e14))) | ~ (h17(op1(e10, e13)) = op2(h17(e10), h17(e13))) | ~ (h17(op1(e10, e12)) = op2(h17(e10), h17(e12))) | ~ (h17(op1(e10, e11)) = op2(h17(e10), h17(e11))) | ~ (h17(op1(e10, e10)) = op2(h17(e10), h17(e10)))) & ((~ (e24 = h16(e14)) & ~ (e24 = h16(e13)) & ~ (e24 = h16(e12)) & ~ (e24 = h16(e11)) & ~ (e24 = h16(e10))) | sP63 | sP62 | sP61 | sP60 | ~ (h16(op1(e14, e14)) = op2(h16(e14), h16(e14))) | ~ (h16(op1(e14, e13)) = op2(h16(e14), h16(e13))) | ~ (h16(op1(e14, e12)) = op2(h16(e14), h16(e12))) | ~ (h16(op1(e14, e11)) = op2(h16(e14), h16(e11))) | ~ (h16(op1(e14, e10)) = op2(h16(e14), h16(e10))) | ~ (h16(op1(e13, e14)) = op2(h16(e13), h16(e14))) | ~ (h16(op1(e13, e13)) = op2(h16(e13), h16(e13))) | ~ (h16(op1(e13, e12)) = op2(h16(e13), h16(e12))) | ~ (h16(op1(e13, e11)) = op2(h16(e13), h16(e11))) | ~ (h16(op1(e13, e10)) = op2(h16(e13), h16(e10))) | ~ (h16(op1(e12, e14)) = op2(h16(e12), h16(e14))) | ~ (h16(op1(e12, e13)) = op2(h16(e12), h16(e13))) | ~ (h16(op1(e12, e12)) = op2(h16(e12), h16(e12))) | ~ (h16(op1(e12, e11)) = op2(h16(e12), h16(e11))) | ~ (h16(op1(e12, e10)) = op2(h16(e12), h16(e10))) | ~ (h16(op1(e11, e14)) = op2(h16(e11), h16(e14))) | ~ (h16(op1(e11, e13)) = op2(h16(e11), h16(e13))) | ~ (h16(op1(e11, e12)) = op2(h16(e11), h16(e12))) | ~ (h16(op1(e11, e11)) = op2(h16(e11), h16(e11))) | ~ (h16(op1(e11, e10)) = op2(h16(e11), h16(e10))) | ~ (h16(op1(e10, e14)) = op2(h16(e10), h16(e14))) | ~ (h16(op1(e10, e13)) = op2(h16(e10), h16(e13))) | ~ (h16(op1(e10, e12)) = op2(h16(e10), h16(e12))) | ~ (h16(op1(e10, e11)) = op2(h16(e10), h16(e11))) | ~ (h16(op1(e10, e10)) = op2(h16(e10), h16(e10)))) & ((~ (e24 = h15(e14)) & ~ (e24 = h15(e13)) & ~ (e24 = h15(e12)) & ~ (e24 = h15(e11)) & ~ (e24 = h15(e10))) | sP59 | sP58 | sP57 | sP56 | ~ (h15(op1(e14, e14)) = op2(h15(e14), h15(e14))) | ~ (h15(op1(e14, e13)) = op2(h15(e14), h15(e13))) | ~ (h15(op1(e14, e12)) = op2(h15(e14), h15(e12))) | ~ (h15(op1(e14, e11)) = op2(h15(e14), h15(e11))) | ~ (h15(op1(e14, e10)) = op2(h15(e14), h15(e10))) | ~ (h15(op1(e13, e14)) = op2(h15(e13), h15(e14))) | ~ (h15(op1(e13, e13)) = op2(h15(e13), h15(e13))) | ~ (h15(op1(e13, e12)) = op2(h15(e13), h15(e12))) | ~ (h15(op1(e13, e11)) = op2(h15(e13), h15(e11))) | ~ (h15(op1(e13, e10)) = op2(h15(e13), h15(e10))) | ~ (h15(op1(e12, e14)) = op2(h15(e12), h15(e14))) | ~ (h15(op1(e12, e13)) = op2(h15(e12), h15(e13))) | ~ (h15(op1(e12, e12)) = op2(h15(e12), h15(e12))) | ~ (h15(op1(e12, e11)) = op2(h15(e12), h15(e11))) | ~ (h15(op1(e12, e10)) = op2(h15(e12), h15(e10))) | ~ (h15(op1(e11, e14)) = op2(h15(e11), h15(e14))) | ~ (h15(op1(e11, e13)) = op2(h15(e11), h15(e13))) | ~ (h15(op1(e11, e12)) = op2(h15(e11), h15(e12))) | ~ (h15(op1(e11, e11)) = op2(h15(e11), h15(e11))) | ~ (h15(op1(e11, e10)) = op2(h15(e11), h15(e10))) | ~ (h15(op1(e10, e14)) = op2(h15(e10), h15(e14))) | ~ (h15(op1(e10, e13)) = op2(h15(e10), h15(e13))) | ~ (h15(op1(e10, e12)) = op2(h15(e10), h15(e12))) | ~ (h15(op1(e10, e11)) = op2(h15(e10), h15(e11))) | ~ (h15(op1(e10, e10)) = op2(h15(e10), h15(e10)))) & ((~ (e24 = h14(e14)) & ~ (e24 = h14(e13)) & ~ (e24 = h14(e12)) & ~ (e24 = h14(e11)) & ~ (e24 = h14(e10))) | sP55 | sP54 | sP53 | sP52 | ~ (h14(op1(e14, e14)) = op2(h14(e14), h14(e14))) | ~ (h14(op1(e14, e13)) = op2(h14(e14), h14(e13))) | ~ (h14(op1(e14, e12)) = op2(h14(e14), h14(e12))) | ~ (h14(op1(e14, e11)) = op2(h14(e14), h14(e11))) | ~ (h14(op1(e14, e10)) = op2(h14(e14), h14(e10))) | ~ (h14(op1(e13, e14)) = op2(h14(e13), h14(e14))) | ~ (h14(op1(e13, e13)) = op2(h14(e13), h14(e13))) | ~ (h14(op1(e13, e12)) = op2(h14(e13), h14(e12))) | ~ (h14(op1(e13, e11)) = op2(h14(e13), h14(e11))) | ~ (h14(op1(e13, e10)) = op2(h14(e13), h14(e10))) | ~ (h14(op1(e12, e14)) = op2(h14(e12), h14(e14))) | ~ (h14(op1(e12, e13)) = op2(h14(e12), h14(e13))) | ~ (h14(op1(e12, e12)) = op2(h14(e12), h14(e12))) | ~ (h14(op1(e12, e11)) = op2(h14(e12), h14(e11))) | ~ (h14(op1(e12, e10)) = op2(h14(e12), h14(e10))) | ~ (h14(op1(e11, e14)) = op2(h14(e11), h14(e14))) | ~ (h14(op1(e11, e13)) = op2(h14(e11), h14(e13))) | ~ (h14(op1(e11, e12)) = op2(h14(e11), h14(e12))) | ~ (h14(op1(e11, e11)) = op2(h14(e11), h14(e11))) | ~ (h14(op1(e11, e10)) = op2(h14(e11), h14(e10))) | ~ (h14(op1(e10, e14)) = op2(h14(e10), h14(e14))) | ~ (h14(op1(e10, e13)) = op2(h14(e10), h14(e13))) | ~ (h14(op1(e10, e12)) = op2(h14(e10), h14(e12))) | ~ (h14(op1(e10, e11)) = op2(h14(e10), h14(e11))) | ~ (h14(op1(e10, e10)) = op2(h14(e10), h14(e10)))) & ((~ (e24 = h13(e14)) & ~ (e24 = h13(e13)) & ~ (e24 = h13(e12)) & ~ (e24 = h13(e11)) & ~ (e24 = h13(e10))) | sP51 | sP50 | sP49 | sP48 | ~ (h13(op1(e14, e14)) = op2(h13(e14), h13(e14))) | ~ (h13(op1(e14, e13)) = op2(h13(e14), h13(e13))) | ~ (h13(op1(e14, e12)) = op2(h13(e14), h13(e12))) | ~ (h13(op1(e14, e11)) = op2(h13(e14), h13(e11))) | ~ (h13(op1(e14, e10)) = op2(h13(e14), h13(e10))) | ~ (h13(op1(e13, e14)) = op2(h13(e13), h13(e14))) | ~ (h13(op1(e13, e13)) = op2(h13(e13), h13(e13))) | ~ (h13(op1(e13, e12)) = op2(h13(e13), h13(e12))) | ~ (h13(op1(e13, e11)) = op2(h13(e13), h13(e11))) | ~ (h13(op1(e13, e10)) = op2(h13(e13), h13(e10))) | ~ (h13(op1(e12, e14)) = op2(h13(e12), h13(e14))) | ~ (h13(op1(e12, e13)) = op2(h13(e12), h13(e13))) | ~ (h13(op1(e12, e12)) = op2(h13(e12), h13(e12))) | ~ (h13(op1(e12, e11)) = op2(h13(e12), h13(e11))) | ~ (h13(op1(e12, e10)) = op2(h13(e12), h13(e10))) | ~ (h13(op1(e11, e14)) = op2(h13(e11), h13(e14))) | ~ (h13(op1(e11, e13)) = op2(h13(e11), h13(e13))) | ~ (h13(op1(e11, e12)) = op2(h13(e11), h13(e12))) | ~ (h13(op1(e11, e11)) = op2(h13(e11), h13(e11))) | ~ (h13(op1(e11, e10)) = op2(h13(e11), h13(e10))) | ~ (h13(op1(e10, e14)) = op2(h13(e10), h13(e14))) | ~ (h13(op1(e10, e13)) = op2(h13(e10), h13(e13))) | ~ (h13(op1(e10, e12)) = op2(h13(e10), h13(e12))) | ~ (h13(op1(e10, e11)) = op2(h13(e10), h13(e11))) | ~ (h13(op1(e10, e10)) = op2(h13(e10), h13(e10)))) & ((~ (e24 = h12(e14)) & ~ (e24 = h12(e13)) & ~ (e24 = h12(e12)) & ~ (e24 = h12(e11)) & ~ (e24 = h12(e10))) | sP47 | sP46 | sP45 | sP44 | ~ (h12(op1(e14, e14)) = op2(h12(e14), h12(e14))) | ~ (h12(op1(e14, e13)) = op2(h12(e14), h12(e13))) | ~ (h12(op1(e14, e12)) = op2(h12(e14), h12(e12))) | ~ (h12(op1(e14, e11)) = op2(h12(e14), h12(e11))) | ~ (h12(op1(e14, e10)) = op2(h12(e14), h12(e10))) | ~ (h12(op1(e13, e14)) = op2(h12(e13), h12(e14))) | ~ (h12(op1(e13, e13)) = op2(h12(e13), h12(e13))) | ~ (h12(op1(e13, e12)) = op2(h12(e13), h12(e12))) | ~ (h12(op1(e13, e11)) = op2(h12(e13), h12(e11))) | ~ (h12(op1(e13, e10)) = op2(h12(e13), h12(e10))) | ~ (h12(op1(e12, e14)) = op2(h12(e12), h12(e14))) | ~ (h12(op1(e12, e13)) = op2(h12(e12), h12(e13))) | ~ (h12(op1(e12, e12)) = op2(h12(e12), h12(e12))) | ~ (h12(op1(e12, e11)) = op2(h12(e12), h12(e11))) | ~ (h12(op1(e12, e10)) = op2(h12(e12), h12(e10))) | ~ (h12(op1(e11, e14)) = op2(h12(e11), h12(e14))) | ~ (h12(op1(e11, e13)) = op2(h12(e11), h12(e13))) | ~ (h12(op1(e11, e12)) = op2(h12(e11), h12(e12))) | ~ (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) | ~ (h12(op1(e11, e10)) = op2(h12(e11), h12(e10))) | ~ (h12(op1(e10, e14)) = op2(h12(e10), h12(e14))) | ~ (h12(op1(e10, e13)) = op2(h12(e10), h12(e13))) | ~ (h12(op1(e10, e12)) = op2(h12(e10), h12(e12))) | ~ (h12(op1(e10, e11)) = op2(h12(e10), h12(e11))) | ~ (h12(op1(e10, e10)) = op2(h12(e10), h12(e10)))) & ((~ (e24 = h11(e14)) & ~ (e24 = h11(e13)) & ~ (e24 = h11(e12)) & ~ (e24 = h11(e11)) & ~ (e24 = h11(e10))) | sP43 | sP42 | sP41 | sP40 | ~ (h11(op1(e14, e14)) = op2(h11(e14), h11(e14))) | ~ (h11(op1(e14, e13)) = op2(h11(e14), h11(e13))) | ~ (h11(op1(e14, e12)) = op2(h11(e14), h11(e12))) | ~ (h11(op1(e14, e11)) = op2(h11(e14), h11(e11))) | ~ (h11(op1(e14, e10)) = op2(h11(e14), h11(e10))) | ~ (h11(op1(e13, e14)) = op2(h11(e13), h11(e14))) | ~ (h11(op1(e13, e13)) = op2(h11(e13), h11(e13))) | ~ (h11(op1(e13, e12)) = op2(h11(e13), h11(e12))) | ~ (h11(op1(e13, e11)) = op2(h11(e13), h11(e11))) | ~ (h11(op1(e13, e10)) = op2(h11(e13), h11(e10))) | ~ (h11(op1(e12, e14)) = op2(h11(e12), h11(e14))) | ~ (h11(op1(e12, e13)) = op2(h11(e12), h11(e13))) | ~ (h11(op1(e12, e12)) = op2(h11(e12), h11(e12))) | ~ (h11(op1(e12, e11)) = op2(h11(e12), h11(e11))) | ~ (h11(op1(e12, e10)) = op2(h11(e12), h11(e10))) | ~ (h11(op1(e11, e14)) = op2(h11(e11), h11(e14))) | ~ (h11(op1(e11, e13)) = op2(h11(e11), h11(e13))) | ~ (h11(op1(e11, e12)) = op2(h11(e11), h11(e12))) | ~ (h11(op1(e11, e11)) = op2(h11(e11), h11(e11))) | ~ (h11(op1(e11, e10)) = op2(h11(e11), h11(e10))) | ~ (h11(op1(e10, e14)) = op2(h11(e10), h11(e14))) | ~ (h11(op1(e10, e13)) = op2(h11(e10), h11(e13))) | ~ (h11(op1(e10, e12)) = op2(h11(e10), h11(e12))) | ~ (h11(op1(e10, e11)) = op2(h11(e10), h11(e11))) | ~ (h11(op1(e10, e10)) = op2(h11(e10), h11(e10)))) & ((~ (e24 = h10(e14)) & ~ (e24 = h10(e13)) & ~ (e24 = h10(e12)) & ~ (e24 = h10(e11)) & ~ (e24 = h10(e10))) | sP39 | sP38 | sP37 | sP36 | ~ (h10(op1(e14, e14)) = op2(h10(e14), h10(e14))) | ~ (h10(op1(e14, e13)) = op2(h10(e14), h10(e13))) | ~ (h10(op1(e14, e12)) = op2(h10(e14), h10(e12))) | ~ (h10(op1(e14, e11)) = op2(h10(e14), h10(e11))) | ~ (h10(op1(e14, e10)) = op2(h10(e14), h10(e10))) | ~ (h10(op1(e13, e14)) = op2(h10(e13), h10(e14))) | ~ (h10(op1(e13, e13)) = op2(h10(e13), h10(e13))) | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) | ~ (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), h10(e14))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) | ~ (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) | ~ (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) | ~ (h10(op1(e11, e14)) = op2(h10(e11), h10(e14))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))) & ((~ (e24 = h9(e14)) & ~ (e24 = h9(e13)) & ~ (e24 = h9(e12)) & ~ (e24 = h9(e11)) & ~ (e24 = h9(e10))) | sP35 | sP34 | sP33 | sP32 | ~ (h9(op1(e14, e14)) = op2(h9(e14), h9(e14))) | ~ (h9(op1(e14, e13)) = op2(h9(e14), h9(e13))) | ~ (h9(op1(e14, e12)) = op2(h9(e14), h9(e12))) | ~ (h9(op1(e14, e11)) = op2(h9(e14), h9(e11))) | ~ (h9(op1(e14, e10)) = op2(h9(e14), h9(e10))) | ~ (h9(op1(e13, e14)) = op2(h9(e13), h9(e14))) | ~ (h9(op1(e13, e13)) = op2(h9(e13), h9(e13))) | ~ (h9(op1(e13, e12)) = op2(h9(e13), h9(e12))) | ~ (h9(op1(e13, e11)) = op2(h9(e13), h9(e11))) | ~ (h9(op1(e13, e10)) = op2(h9(e13), h9(e10))) | ~ (h9(op1(e12, e14)) = op2(h9(e12), h9(e14))) | ~ (h9(op1(e12, e13)) = op2(h9(e12), h9(e13))) | ~ (h9(op1(e12, e12)) = op2(h9(e12), h9(e12))) | ~ (h9(op1(e12, e11)) = op2(h9(e12), h9(e11))) | ~ (h9(op1(e12, e10)) = op2(h9(e12), h9(e10))) | ~ (h9(op1(e11, e14)) = op2(h9(e11), h9(e14))) | ~ (h9(op1(e11, e13)) = op2(h9(e11), h9(e13))) | ~ (h9(op1(e11, e12)) = op2(h9(e11), h9(e12))) | ~ (h9(op1(e11, e11)) = op2(h9(e11), h9(e11))) | ~ (h9(op1(e11, e10)) = op2(h9(e11), h9(e10))) | ~ (h9(op1(e10, e14)) = op2(h9(e10), h9(e14))) | ~ (h9(op1(e10, e13)) = op2(h9(e10), h9(e13))) | ~ (h9(op1(e10, e12)) = op2(h9(e10), h9(e12))) | ~ (h9(op1(e10, e11)) = op2(h9(e10), h9(e11))) | ~ (h9(op1(e10, e10)) = op2(h9(e10), h9(e10)))) & ((~ (e24 = h8(e14)) & ~ (e24 = h8(e13)) & ~ (e24 = h8(e12)) & ~ (e24 = h8(e11)) & ~ (e24 = h8(e10))) | sP31 | sP30 | sP29 | sP28 | ~ (h8(op1(e14, e14)) = op2(h8(e14), h8(e14))) | ~ (h8(op1(e14, e13)) = op2(h8(e14), h8(e13))) | ~ (h8(op1(e14, e12)) = op2(h8(e14), h8(e12))) | ~ (h8(op1(e14, e11)) = op2(h8(e14), h8(e11))) | ~ (h8(op1(e14, e10)) = op2(h8(e14), h8(e10))) | ~ (h8(op1(e13, e14)) = op2(h8(e13), h8(e14))) | ~ (h8(op1(e13, e13)) = op2(h8(e13), h8(e13))) | ~ (h8(op1(e13, e12)) = op2(h8(e13), h8(e12))) | ~ (h8(op1(e13, e11)) = op2(h8(e13), h8(e11))) | ~ (h8(op1(e13, e10)) = op2(h8(e13), h8(e10))) | ~ (h8(op1(e12, e14)) = op2(h8(e12), h8(e14))) | ~ (h8(op1(e12, e13)) = op2(h8(e12), h8(e13))) | ~ (h8(op1(e12, e12)) = op2(h8(e12), h8(e12))) | ~ (h8(op1(e12, e11)) = op2(h8(e12), h8(e11))) | ~ (h8(op1(e12, e10)) = op2(h8(e12), h8(e10))) | ~ (h8(op1(e11, e14)) = op2(h8(e11), h8(e14))) | ~ (h8(op1(e11, e13)) = op2(h8(e11), h8(e13))) | ~ (h8(op1(e11, e12)) = op2(h8(e11), h8(e12))) | ~ (h8(op1(e11, e11)) = op2(h8(e11), h8(e11))) | ~ (h8(op1(e11, e10)) = op2(h8(e11), h8(e10))) | ~ (h8(op1(e10, e14)) = op2(h8(e10), h8(e14))) | ~ (h8(op1(e10, e13)) = op2(h8(e10), h8(e13))) | ~ (h8(op1(e10, e12)) = op2(h8(e10), h8(e12))) | ~ (h8(op1(e10, e11)) = op2(h8(e10), h8(e11))) | ~ (h8(op1(e10, e10)) = op2(h8(e10), h8(e10)))) & ((~ (e24 = h7(e14)) & ~ (e24 = h7(e13)) & ~ (e24 = h7(e12)) & ~ (e24 = h7(e11)) & ~ (e24 = h7(e10))) | sP27 | sP26 | sP25 | sP24 | ~ (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) | ~ (h7(op1(e14, e13)) = op2(h7(e14), h7(e13))) | ~ (h7(op1(e14, e12)) = op2(h7(e14), h7(e12))) | ~ (h7(op1(e14, e11)) = op2(h7(e14), h7(e11))) | ~ (h7(op1(e14, e10)) = op2(h7(e14), h7(e10))) | ~ (h7(op1(e13, e14)) = op2(h7(e13), h7(e14))) | ~ (h7(op1(e13, e13)) = op2(h7(e13), h7(e13))) | ~ (h7(op1(e13, e12)) = op2(h7(e13), h7(e12))) | ~ (h7(op1(e13, e11)) = op2(h7(e13), h7(e11))) | ~ (h7(op1(e13, e10)) = op2(h7(e13), h7(e10))) | ~ (h7(op1(e12, e14)) = op2(h7(e12), h7(e14))) | ~ (h7(op1(e12, e13)) = op2(h7(e12), h7(e13))) | ~ (h7(op1(e12, e12)) = op2(h7(e12), h7(e12))) | ~ (h7(op1(e12, e11)) = op2(h7(e12), h7(e11))) | ~ (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) | ~ (h7(op1(e11, e14)) = op2(h7(e11), h7(e14))) | ~ (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) | ~ (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) | ~ (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) | ~ (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) | ~ (h7(op1(e10, e14)) = op2(h7(e10), h7(e14))) | ~ (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) | ~ (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) | ~ (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))) & ((~ (e24 = h6(e14)) & ~ (e24 = h6(e13)) & ~ (e24 = h6(e12)) & ~ (e24 = h6(e11)) & ~ (e24 = h6(e10))) | sP23 | sP22 | sP21 | sP20 | ~ (h6(op1(e14, e14)) = op2(h6(e14), h6(e14))) | ~ (h6(op1(e14, e13)) = op2(h6(e14), h6(e13))) | ~ (h6(op1(e14, e12)) = op2(h6(e14), h6(e12))) | ~ (h6(op1(e14, e11)) = op2(h6(e14), h6(e11))) | ~ (h6(op1(e14, e10)) = op2(h6(e14), h6(e10))) | ~ (h6(op1(e13, e14)) = op2(h6(e13), h6(e14))) | ~ (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) | ~ (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))) & ((~ (e24 = h5(e14)) & ~ (e24 = h5(e13)) & ~ (e24 = h5(e12)) & ~ (e24 = h5(e11)) & ~ (e24 = h5(e10))) | sP19 | sP18 | sP17 | sP16 | ~ (h5(op1(e14, e14)) = op2(h5(e14), h5(e14))) | ~ (h5(op1(e14, e13)) = op2(h5(e14), h5(e13))) | ~ (h5(op1(e14, e12)) = op2(h5(e14), h5(e12))) | ~ (h5(op1(e14, e11)) = op2(h5(e14), h5(e11))) | ~ (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) | ~ (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))) & ((~ (e24 = h4(e14)) & ~ (e24 = h4(e13)) & ~ (e24 = h4(e12)) & ~ (e24 = h4(e11)) & ~ (e24 = h4(e10))) | sP15 | sP14 | sP13 | sP12 | ~ (h4(op1(e14, e14)) = op2(h4(e14), h4(e14))) | ~ (h4(op1(e14, e13)) = op2(h4(e14), h4(e13))) | ~ (h4(op1(e14, e12)) = op2(h4(e14), h4(e12))) | ~ (h4(op1(e14, e11)) = op2(h4(e14), h4(e11))) | ~ (h4(op1(e14, e10)) = op2(h4(e14), h4(e10))) | ~ (h4(op1(e13, e14)) = op2(h4(e13), h4(e14))) | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e14)) = op2(h4(e12), h4(e14))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e14)) = op2(h4(e11), h4(e14))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e14)) = op2(h4(e10), h4(e14))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) & ((~ (e24 = h3(e14)) & ~ (e24 = h3(e13)) & ~ (e24 = h3(e12)) & ~ (e24 = h3(e11)) & ~ (e24 = h3(e10))) | sP11 | sP10 | sP9 | sP8 | ~ (h3(op1(e14, e14)) = op2(h3(e14), h3(e14))) | ~ (h3(op1(e14, e13)) = op2(h3(e14), h3(e13))) | ~ (h3(op1(e14, e12)) = op2(h3(e14), h3(e12))) | ~ (h3(op1(e14, e11)) = op2(h3(e14), h3(e11))) | ~ (h3(op1(e14, e10)) = op2(h3(e14), h3(e10))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), h3(e14))) | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) & ((~ (e24 = h2(e14)) & ~ (e24 = h2(e13)) & ~ (e24 = h2(e12)) & ~ (e24 = h2(e11)) & ~ (e24 = h2(e10))) | sP7 | sP6 | sP5 | sP4 | ~ (h2(op1(e14, e14)) = op2(h2(e14), h2(e14))) | ~ (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) | ~ (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) | ~ (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) | ~ (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) | ~ (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) & ((~ (e24 = h1(e14)) & ~ (e24 = h1(e13)) & ~ (e24 = h1(e12)) & ~ (e24 = h1(e11)) & ~ (e24 = h1(e10))) | sP3 | sP2 | sP1 | sP0 | ~ (h1(op1(e14, e14)) = op2(h1(e14), h1(e14))) | ~ (h1(op1(e14, e13)) = op2(h1(e14), h1(e13))) | ~ (h1(op1(e14, e12)) = op2(h1(e14), h1(e12))) | ~ (h1(op1(e14, e11)) = op2(h1(e14), h1(e11))) | ~ (h1(op1(e14, e10)) = op2(h1(e14), h1(e10))) | ~ (h1(op1(e13, e14)) = op2(h1(e13), h1(e14))) | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e14)) = op2(h1(e12), h1(e14))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e14)) = op2(h1(e11), h1(e14))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e14)) = op2(h1(e10), h1(e14))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(definition_folding, [], [f38, e118, e117, e116, e115, e114, e113, e112, e111, e110, e109, e108, e107, e106, e105, e104, e103, e102, e101, e100, e99, e98, e97, e96, e95, e94, e93, e92, e91, e90, e89, e88, e87, e86, e85, e84, e83, e82, e81, e80, e79, e78, e77, e76, e75, e74, e73, e72, e71, e70, e69, e68, e67, e66, e65, e64, e63, e62, e61, e60, e59, e58, e57, e56, e55, e54, e53, e52, e51, e50, e49, e48, e47, e46, e45, e44, e43, e42, e41, e40, e39])).
fof(f39, plain, ((~ (e20 = h1(e14)) & ~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10))) | ~ sP0), inference(usedef, [], [e39])).
fof(e39, plain, (sP0 <=> (~ (e20 = h1(e14)) & ~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f40, plain, ((~ (e21 = h1(e14)) & ~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10))) | ~ sP1), inference(usedef, [], [e40])).
fof(e40, plain, (sP1 <=> (~ (e21 = h1(e14)) & ~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f41, plain, ((~ (e22 = h1(e14)) & ~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10))) | ~ sP2), inference(usedef, [], [e41])).
fof(e41, plain, (sP2 <=> (~ (e22 = h1(e14)) & ~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f42, plain, ((~ (e23 = h1(e14)) & ~ (e23 = h1(e13)) & ~ (e23 = h1(e12)) & ~ (e23 = h1(e11)) & ~ (e23 = h1(e10))) | ~ sP3), inference(usedef, [], [e42])).
fof(e42, plain, (sP3 <=> (~ (e23 = h1(e14)) & ~ (e23 = h1(e13)) & ~ (e23 = h1(e12)) & ~ (e23 = h1(e11)) & ~ (e23 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f43, plain, ((~ (e20 = h2(e14)) & ~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10))) | ~ sP4), inference(usedef, [], [e43])).
fof(e43, plain, (sP4 <=> (~ (e20 = h2(e14)) & ~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f44, plain, ((~ (e21 = h2(e14)) & ~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10))) | ~ sP5), inference(usedef, [], [e44])).
fof(e44, plain, (sP5 <=> (~ (e21 = h2(e14)) & ~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f45, plain, ((~ (e22 = h2(e14)) & ~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10))) | ~ sP6), inference(usedef, [], [e45])).
fof(e45, plain, (sP6 <=> (~ (e22 = h2(e14)) & ~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f46, plain, ((~ (e23 = h2(e14)) & ~ (e23 = h2(e13)) & ~ (e23 = h2(e12)) & ~ (e23 = h2(e11)) & ~ (e23 = h2(e10))) | ~ sP7), inference(usedef, [], [e46])).
fof(e46, plain, (sP7 <=> (~ (e23 = h2(e14)) & ~ (e23 = h2(e13)) & ~ (e23 = h2(e12)) & ~ (e23 = h2(e11)) & ~ (e23 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f47, plain, ((~ (e20 = h3(e14)) & ~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10))) | ~ sP8), inference(usedef, [], [e47])).
fof(e47, plain, (sP8 <=> (~ (e20 = h3(e14)) & ~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f48, plain, ((~ (e21 = h3(e14)) & ~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10))) | ~ sP9), inference(usedef, [], [e48])).
fof(e48, plain, (sP9 <=> (~ (e21 = h3(e14)) & ~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f49, plain, ((~ (e22 = h3(e14)) & ~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10))) | ~ sP10), inference(usedef, [], [e49])).
fof(e49, plain, (sP10 <=> (~ (e22 = h3(e14)) & ~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f50, plain, ((~ (e23 = h3(e14)) & ~ (e23 = h3(e13)) & ~ (e23 = h3(e12)) & ~ (e23 = h3(e11)) & ~ (e23 = h3(e10))) | ~ sP11), inference(usedef, [], [e50])).
fof(e50, plain, (sP11 <=> (~ (e23 = h3(e14)) & ~ (e23 = h3(e13)) & ~ (e23 = h3(e12)) & ~ (e23 = h3(e11)) & ~ (e23 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f51, plain, ((~ (e20 = h4(e14)) & ~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10))) | ~ sP12), inference(usedef, [], [e51])).
fof(e51, plain, (sP12 <=> (~ (e20 = h4(e14)) & ~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f52, plain, ((~ (e21 = h4(e14)) & ~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10))) | ~ sP13), inference(usedef, [], [e52])).
fof(e52, plain, (sP13 <=> (~ (e21 = h4(e14)) & ~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f53, plain, ((~ (e22 = h4(e14)) & ~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10))) | ~ sP14), inference(usedef, [], [e53])).
fof(e53, plain, (sP14 <=> (~ (e22 = h4(e14)) & ~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f54, plain, ((~ (e23 = h4(e14)) & ~ (e23 = h4(e13)) & ~ (e23 = h4(e12)) & ~ (e23 = h4(e11)) & ~ (e23 = h4(e10))) | ~ sP15), inference(usedef, [], [e54])).
fof(e54, plain, (sP15 <=> (~ (e23 = h4(e14)) & ~ (e23 = h4(e13)) & ~ (e23 = h4(e12)) & ~ (e23 = h4(e11)) & ~ (e23 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f55, plain, ((~ (e20 = h5(e14)) & ~ (e20 = h5(e13)) & ~ (e20 = h5(e12)) & ~ (e20 = h5(e11)) & ~ (e20 = h5(e10))) | ~ sP16), inference(usedef, [], [e55])).
fof(e55, plain, (sP16 <=> (~ (e20 = h5(e14)) & ~ (e20 = h5(e13)) & ~ (e20 = h5(e12)) & ~ (e20 = h5(e11)) & ~ (e20 = h5(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f56, plain, ((~ (e21 = h5(e14)) & ~ (e21 = h5(e13)) & ~ (e21 = h5(e12)) & ~ (e21 = h5(e11)) & ~ (e21 = h5(e10))) | ~ sP17), inference(usedef, [], [e56])).
fof(e56, plain, (sP17 <=> (~ (e21 = h5(e14)) & ~ (e21 = h5(e13)) & ~ (e21 = h5(e12)) & ~ (e21 = h5(e11)) & ~ (e21 = h5(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f57, plain, ((~ (e22 = h5(e14)) & ~ (e22 = h5(e13)) & ~ (e22 = h5(e12)) & ~ (e22 = h5(e11)) & ~ (e22 = h5(e10))) | ~ sP18), inference(usedef, [], [e57])).
fof(e57, plain, (sP18 <=> (~ (e22 = h5(e14)) & ~ (e22 = h5(e13)) & ~ (e22 = h5(e12)) & ~ (e22 = h5(e11)) & ~ (e22 = h5(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f58, plain, ((~ (e23 = h5(e14)) & ~ (e23 = h5(e13)) & ~ (e23 = h5(e12)) & ~ (e23 = h5(e11)) & ~ (e23 = h5(e10))) | ~ sP19), inference(usedef, [], [e58])).
fof(e58, plain, (sP19 <=> (~ (e23 = h5(e14)) & ~ (e23 = h5(e13)) & ~ (e23 = h5(e12)) & ~ (e23 = h5(e11)) & ~ (e23 = h5(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP19])])).
fof(f59, plain, ((~ (e20 = h6(e14)) & ~ (e20 = h6(e13)) & ~ (e20 = h6(e12)) & ~ (e20 = h6(e11)) & ~ (e20 = h6(e10))) | ~ sP20), inference(usedef, [], [e59])).
fof(e59, plain, (sP20 <=> (~ (e20 = h6(e14)) & ~ (e20 = h6(e13)) & ~ (e20 = h6(e12)) & ~ (e20 = h6(e11)) & ~ (e20 = h6(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP20])])).
fof(f60, plain, ((~ (e21 = h6(e14)) & ~ (e21 = h6(e13)) & ~ (e21 = h6(e12)) & ~ (e21 = h6(e11)) & ~ (e21 = h6(e10))) | ~ sP21), inference(usedef, [], [e60])).
fof(e60, plain, (sP21 <=> (~ (e21 = h6(e14)) & ~ (e21 = h6(e13)) & ~ (e21 = h6(e12)) & ~ (e21 = h6(e11)) & ~ (e21 = h6(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP21])])).
fof(f61, plain, ((~ (e22 = h6(e14)) & ~ (e22 = h6(e13)) & ~ (e22 = h6(e12)) & ~ (e22 = h6(e11)) & ~ (e22 = h6(e10))) | ~ sP22), inference(usedef, [], [e61])).
fof(e61, plain, (sP22 <=> (~ (e22 = h6(e14)) & ~ (e22 = h6(e13)) & ~ (e22 = h6(e12)) & ~ (e22 = h6(e11)) & ~ (e22 = h6(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP22])])).
fof(f62, plain, ((~ (e23 = h6(e14)) & ~ (e23 = h6(e13)) & ~ (e23 = h6(e12)) & ~ (e23 = h6(e11)) & ~ (e23 = h6(e10))) | ~ sP23), inference(usedef, [], [e62])).
fof(e62, plain, (sP23 <=> (~ (e23 = h6(e14)) & ~ (e23 = h6(e13)) & ~ (e23 = h6(e12)) & ~ (e23 = h6(e11)) & ~ (e23 = h6(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP23])])).
fof(f63, plain, ((~ (e20 = h7(e14)) & ~ (e20 = h7(e13)) & ~ (e20 = h7(e12)) & ~ (e20 = h7(e11)) & ~ (e20 = h7(e10))) | ~ sP24), inference(usedef, [], [e63])).
fof(e63, plain, (sP24 <=> (~ (e20 = h7(e14)) & ~ (e20 = h7(e13)) & ~ (e20 = h7(e12)) & ~ (e20 = h7(e11)) & ~ (e20 = h7(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP24])])).
fof(f64, plain, ((~ (e21 = h7(e14)) & ~ (e21 = h7(e13)) & ~ (e21 = h7(e12)) & ~ (e21 = h7(e11)) & ~ (e21 = h7(e10))) | ~ sP25), inference(usedef, [], [e64])).
fof(e64, plain, (sP25 <=> (~ (e21 = h7(e14)) & ~ (e21 = h7(e13)) & ~ (e21 = h7(e12)) & ~ (e21 = h7(e11)) & ~ (e21 = h7(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP25])])).
fof(f65, plain, ((~ (e22 = h7(e14)) & ~ (e22 = h7(e13)) & ~ (e22 = h7(e12)) & ~ (e22 = h7(e11)) & ~ (e22 = h7(e10))) | ~ sP26), inference(usedef, [], [e65])).
fof(e65, plain, (sP26 <=> (~ (e22 = h7(e14)) & ~ (e22 = h7(e13)) & ~ (e22 = h7(e12)) & ~ (e22 = h7(e11)) & ~ (e22 = h7(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP26])])).
fof(f66, plain, ((~ (e23 = h7(e14)) & ~ (e23 = h7(e13)) & ~ (e23 = h7(e12)) & ~ (e23 = h7(e11)) & ~ (e23 = h7(e10))) | ~ sP27), inference(usedef, [], [e66])).
fof(e66, plain, (sP27 <=> (~ (e23 = h7(e14)) & ~ (e23 = h7(e13)) & ~ (e23 = h7(e12)) & ~ (e23 = h7(e11)) & ~ (e23 = h7(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP27])])).
fof(f67, plain, ((~ (e20 = h8(e14)) & ~ (e20 = h8(e13)) & ~ (e20 = h8(e12)) & ~ (e20 = h8(e11)) & ~ (e20 = h8(e10))) | ~ sP28), inference(usedef, [], [e67])).
fof(e67, plain, (sP28 <=> (~ (e20 = h8(e14)) & ~ (e20 = h8(e13)) & ~ (e20 = h8(e12)) & ~ (e20 = h8(e11)) & ~ (e20 = h8(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP28])])).
fof(f68, plain, ((~ (e21 = h8(e14)) & ~ (e21 = h8(e13)) & ~ (e21 = h8(e12)) & ~ (e21 = h8(e11)) & ~ (e21 = h8(e10))) | ~ sP29), inference(usedef, [], [e68])).
fof(e68, plain, (sP29 <=> (~ (e21 = h8(e14)) & ~ (e21 = h8(e13)) & ~ (e21 = h8(e12)) & ~ (e21 = h8(e11)) & ~ (e21 = h8(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP29])])).
fof(f69, plain, ((~ (e22 = h8(e14)) & ~ (e22 = h8(e13)) & ~ (e22 = h8(e12)) & ~ (e22 = h8(e11)) & ~ (e22 = h8(e10))) | ~ sP30), inference(usedef, [], [e69])).
fof(e69, plain, (sP30 <=> (~ (e22 = h8(e14)) & ~ (e22 = h8(e13)) & ~ (e22 = h8(e12)) & ~ (e22 = h8(e11)) & ~ (e22 = h8(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP30])])).
fof(f70, plain, ((~ (e23 = h8(e14)) & ~ (e23 = h8(e13)) & ~ (e23 = h8(e12)) & ~ (e23 = h8(e11)) & ~ (e23 = h8(e10))) | ~ sP31), inference(usedef, [], [e70])).
fof(e70, plain, (sP31 <=> (~ (e23 = h8(e14)) & ~ (e23 = h8(e13)) & ~ (e23 = h8(e12)) & ~ (e23 = h8(e11)) & ~ (e23 = h8(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP31])])).
fof(f71, plain, ((~ (e20 = h9(e14)) & ~ (e20 = h9(e13)) & ~ (e20 = h9(e12)) & ~ (e20 = h9(e11)) & ~ (e20 = h9(e10))) | ~ sP32), inference(usedef, [], [e71])).
fof(e71, plain, (sP32 <=> (~ (e20 = h9(e14)) & ~ (e20 = h9(e13)) & ~ (e20 = h9(e12)) & ~ (e20 = h9(e11)) & ~ (e20 = h9(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP32])])).
fof(f72, plain, ((~ (e21 = h9(e14)) & ~ (e21 = h9(e13)) & ~ (e21 = h9(e12)) & ~ (e21 = h9(e11)) & ~ (e21 = h9(e10))) | ~ sP33), inference(usedef, [], [e72])).
fof(e72, plain, (sP33 <=> (~ (e21 = h9(e14)) & ~ (e21 = h9(e13)) & ~ (e21 = h9(e12)) & ~ (e21 = h9(e11)) & ~ (e21 = h9(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP33])])).
fof(f73, plain, ((~ (e22 = h9(e14)) & ~ (e22 = h9(e13)) & ~ (e22 = h9(e12)) & ~ (e22 = h9(e11)) & ~ (e22 = h9(e10))) | ~ sP34), inference(usedef, [], [e73])).
fof(e73, plain, (sP34 <=> (~ (e22 = h9(e14)) & ~ (e22 = h9(e13)) & ~ (e22 = h9(e12)) & ~ (e22 = h9(e11)) & ~ (e22 = h9(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP34])])).
fof(f74, plain, ((~ (e23 = h9(e14)) & ~ (e23 = h9(e13)) & ~ (e23 = h9(e12)) & ~ (e23 = h9(e11)) & ~ (e23 = h9(e10))) | ~ sP35), inference(usedef, [], [e74])).
fof(e74, plain, (sP35 <=> (~ (e23 = h9(e14)) & ~ (e23 = h9(e13)) & ~ (e23 = h9(e12)) & ~ (e23 = h9(e11)) & ~ (e23 = h9(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP35])])).
fof(f75, plain, ((~ (e20 = h10(e14)) & ~ (e20 = h10(e13)) & ~ (e20 = h10(e12)) & ~ (e20 = h10(e11)) & ~ (e20 = h10(e10))) | ~ sP36), inference(usedef, [], [e75])).
fof(e75, plain, (sP36 <=> (~ (e20 = h10(e14)) & ~ (e20 = h10(e13)) & ~ (e20 = h10(e12)) & ~ (e20 = h10(e11)) & ~ (e20 = h10(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP36])])).
fof(f76, plain, ((~ (e21 = h10(e14)) & ~ (e21 = h10(e13)) & ~ (e21 = h10(e12)) & ~ (e21 = h10(e11)) & ~ (e21 = h10(e10))) | ~ sP37), inference(usedef, [], [e76])).
fof(e76, plain, (sP37 <=> (~ (e21 = h10(e14)) & ~ (e21 = h10(e13)) & ~ (e21 = h10(e12)) & ~ (e21 = h10(e11)) & ~ (e21 = h10(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP37])])).
fof(f77, plain, ((~ (e22 = h10(e14)) & ~ (e22 = h10(e13)) & ~ (e22 = h10(e12)) & ~ (e22 = h10(e11)) & ~ (e22 = h10(e10))) | ~ sP38), inference(usedef, [], [e77])).
fof(e77, plain, (sP38 <=> (~ (e22 = h10(e14)) & ~ (e22 = h10(e13)) & ~ (e22 = h10(e12)) & ~ (e22 = h10(e11)) & ~ (e22 = h10(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP38])])).
fof(f78, plain, ((~ (e23 = h10(e14)) & ~ (e23 = h10(e13)) & ~ (e23 = h10(e12)) & ~ (e23 = h10(e11)) & ~ (e23 = h10(e10))) | ~ sP39), inference(usedef, [], [e78])).
fof(e78, plain, (sP39 <=> (~ (e23 = h10(e14)) & ~ (e23 = h10(e13)) & ~ (e23 = h10(e12)) & ~ (e23 = h10(e11)) & ~ (e23 = h10(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP39])])).
fof(f79, plain, ((~ (e20 = h11(e14)) & ~ (e20 = h11(e13)) & ~ (e20 = h11(e12)) & ~ (e20 = h11(e11)) & ~ (e20 = h11(e10))) | ~ sP40), inference(usedef, [], [e79])).
fof(e79, plain, (sP40 <=> (~ (e20 = h11(e14)) & ~ (e20 = h11(e13)) & ~ (e20 = h11(e12)) & ~ (e20 = h11(e11)) & ~ (e20 = h11(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP40])])).
fof(f80, plain, ((~ (e21 = h11(e14)) & ~ (e21 = h11(e13)) & ~ (e21 = h11(e12)) & ~ (e21 = h11(e11)) & ~ (e21 = h11(e10))) | ~ sP41), inference(usedef, [], [e80])).
fof(e80, plain, (sP41 <=> (~ (e21 = h11(e14)) & ~ (e21 = h11(e13)) & ~ (e21 = h11(e12)) & ~ (e21 = h11(e11)) & ~ (e21 = h11(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP41])])).
fof(f81, plain, ((~ (e22 = h11(e14)) & ~ (e22 = h11(e13)) & ~ (e22 = h11(e12)) & ~ (e22 = h11(e11)) & ~ (e22 = h11(e10))) | ~ sP42), inference(usedef, [], [e81])).
fof(e81, plain, (sP42 <=> (~ (e22 = h11(e14)) & ~ (e22 = h11(e13)) & ~ (e22 = h11(e12)) & ~ (e22 = h11(e11)) & ~ (e22 = h11(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP42])])).
fof(f82, plain, ((~ (e23 = h11(e14)) & ~ (e23 = h11(e13)) & ~ (e23 = h11(e12)) & ~ (e23 = h11(e11)) & ~ (e23 = h11(e10))) | ~ sP43), inference(usedef, [], [e82])).
fof(e82, plain, (sP43 <=> (~ (e23 = h11(e14)) & ~ (e23 = h11(e13)) & ~ (e23 = h11(e12)) & ~ (e23 = h11(e11)) & ~ (e23 = h11(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP43])])).
fof(f83, plain, ((~ (e20 = h12(e14)) & ~ (e20 = h12(e13)) & ~ (e20 = h12(e12)) & ~ (e20 = h12(e11)) & ~ (e20 = h12(e10))) | ~ sP44), inference(usedef, [], [e83])).
fof(e83, plain, (sP44 <=> (~ (e20 = h12(e14)) & ~ (e20 = h12(e13)) & ~ (e20 = h12(e12)) & ~ (e20 = h12(e11)) & ~ (e20 = h12(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP44])])).
fof(f84, plain, ((~ (e21 = h12(e14)) & ~ (e21 = h12(e13)) & ~ (e21 = h12(e12)) & ~ (e21 = h12(e11)) & ~ (e21 = h12(e10))) | ~ sP45), inference(usedef, [], [e84])).
fof(e84, plain, (sP45 <=> (~ (e21 = h12(e14)) & ~ (e21 = h12(e13)) & ~ (e21 = h12(e12)) & ~ (e21 = h12(e11)) & ~ (e21 = h12(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP45])])).
fof(f85, plain, ((~ (e22 = h12(e14)) & ~ (e22 = h12(e13)) & ~ (e22 = h12(e12)) & ~ (e22 = h12(e11)) & ~ (e22 = h12(e10))) | ~ sP46), inference(usedef, [], [e85])).
fof(e85, plain, (sP46 <=> (~ (e22 = h12(e14)) & ~ (e22 = h12(e13)) & ~ (e22 = h12(e12)) & ~ (e22 = h12(e11)) & ~ (e22 = h12(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP46])])).
fof(f86, plain, ((~ (e23 = h12(e14)) & ~ (e23 = h12(e13)) & ~ (e23 = h12(e12)) & ~ (e23 = h12(e11)) & ~ (e23 = h12(e10))) | ~ sP47), inference(usedef, [], [e86])).
fof(e86, plain, (sP47 <=> (~ (e23 = h12(e14)) & ~ (e23 = h12(e13)) & ~ (e23 = h12(e12)) & ~ (e23 = h12(e11)) & ~ (e23 = h12(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP47])])).
fof(f87, plain, ((~ (e20 = h13(e14)) & ~ (e20 = h13(e13)) & ~ (e20 = h13(e12)) & ~ (e20 = h13(e11)) & ~ (e20 = h13(e10))) | ~ sP48), inference(usedef, [], [e87])).
fof(e87, plain, (sP48 <=> (~ (e20 = h13(e14)) & ~ (e20 = h13(e13)) & ~ (e20 = h13(e12)) & ~ (e20 = h13(e11)) & ~ (e20 = h13(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP48])])).
fof(f88, plain, ((~ (e21 = h13(e14)) & ~ (e21 = h13(e13)) & ~ (e21 = h13(e12)) & ~ (e21 = h13(e11)) & ~ (e21 = h13(e10))) | ~ sP49), inference(usedef, [], [e88])).
fof(e88, plain, (sP49 <=> (~ (e21 = h13(e14)) & ~ (e21 = h13(e13)) & ~ (e21 = h13(e12)) & ~ (e21 = h13(e11)) & ~ (e21 = h13(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP49])])).
fof(f89, plain, ((~ (e22 = h13(e14)) & ~ (e22 = h13(e13)) & ~ (e22 = h13(e12)) & ~ (e22 = h13(e11)) & ~ (e22 = h13(e10))) | ~ sP50), inference(usedef, [], [e89])).
fof(e89, plain, (sP50 <=> (~ (e22 = h13(e14)) & ~ (e22 = h13(e13)) & ~ (e22 = h13(e12)) & ~ (e22 = h13(e11)) & ~ (e22 = h13(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP50])])).
fof(f90, plain, ((~ (e23 = h13(e14)) & ~ (e23 = h13(e13)) & ~ (e23 = h13(e12)) & ~ (e23 = h13(e11)) & ~ (e23 = h13(e10))) | ~ sP51), inference(usedef, [], [e90])).
fof(e90, plain, (sP51 <=> (~ (e23 = h13(e14)) & ~ (e23 = h13(e13)) & ~ (e23 = h13(e12)) & ~ (e23 = h13(e11)) & ~ (e23 = h13(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP51])])).
fof(f91, plain, ((~ (e20 = h14(e14)) & ~ (e20 = h14(e13)) & ~ (e20 = h14(e12)) & ~ (e20 = h14(e11)) & ~ (e20 = h14(e10))) | ~ sP52), inference(usedef, [], [e91])).
fof(e91, plain, (sP52 <=> (~ (e20 = h14(e14)) & ~ (e20 = h14(e13)) & ~ (e20 = h14(e12)) & ~ (e20 = h14(e11)) & ~ (e20 = h14(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP52])])).
fof(f92, plain, ((~ (e21 = h14(e14)) & ~ (e21 = h14(e13)) & ~ (e21 = h14(e12)) & ~ (e21 = h14(e11)) & ~ (e21 = h14(e10))) | ~ sP53), inference(usedef, [], [e92])).
fof(e92, plain, (sP53 <=> (~ (e21 = h14(e14)) & ~ (e21 = h14(e13)) & ~ (e21 = h14(e12)) & ~ (e21 = h14(e11)) & ~ (e21 = h14(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP53])])).
fof(f93, plain, ((~ (e22 = h14(e14)) & ~ (e22 = h14(e13)) & ~ (e22 = h14(e12)) & ~ (e22 = h14(e11)) & ~ (e22 = h14(e10))) | ~ sP54), inference(usedef, [], [e93])).
fof(e93, plain, (sP54 <=> (~ (e22 = h14(e14)) & ~ (e22 = h14(e13)) & ~ (e22 = h14(e12)) & ~ (e22 = h14(e11)) & ~ (e22 = h14(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP54])])).
fof(f94, plain, ((~ (e23 = h14(e14)) & ~ (e23 = h14(e13)) & ~ (e23 = h14(e12)) & ~ (e23 = h14(e11)) & ~ (e23 = h14(e10))) | ~ sP55), inference(usedef, [], [e94])).
fof(e94, plain, (sP55 <=> (~ (e23 = h14(e14)) & ~ (e23 = h14(e13)) & ~ (e23 = h14(e12)) & ~ (e23 = h14(e11)) & ~ (e23 = h14(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP55])])).
fof(f95, plain, ((~ (e20 = h15(e14)) & ~ (e20 = h15(e13)) & ~ (e20 = h15(e12)) & ~ (e20 = h15(e11)) & ~ (e20 = h15(e10))) | ~ sP56), inference(usedef, [], [e95])).
fof(e95, plain, (sP56 <=> (~ (e20 = h15(e14)) & ~ (e20 = h15(e13)) & ~ (e20 = h15(e12)) & ~ (e20 = h15(e11)) & ~ (e20 = h15(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP56])])).
fof(f96, plain, ((~ (e21 = h15(e14)) & ~ (e21 = h15(e13)) & ~ (e21 = h15(e12)) & ~ (e21 = h15(e11)) & ~ (e21 = h15(e10))) | ~ sP57), inference(usedef, [], [e96])).
fof(e96, plain, (sP57 <=> (~ (e21 = h15(e14)) & ~ (e21 = h15(e13)) & ~ (e21 = h15(e12)) & ~ (e21 = h15(e11)) & ~ (e21 = h15(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP57])])).
fof(f97, plain, ((~ (e22 = h15(e14)) & ~ (e22 = h15(e13)) & ~ (e22 = h15(e12)) & ~ (e22 = h15(e11)) & ~ (e22 = h15(e10))) | ~ sP58), inference(usedef, [], [e97])).
fof(e97, plain, (sP58 <=> (~ (e22 = h15(e14)) & ~ (e22 = h15(e13)) & ~ (e22 = h15(e12)) & ~ (e22 = h15(e11)) & ~ (e22 = h15(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP58])])).
fof(f98, plain, ((~ (e23 = h15(e14)) & ~ (e23 = h15(e13)) & ~ (e23 = h15(e12)) & ~ (e23 = h15(e11)) & ~ (e23 = h15(e10))) | ~ sP59), inference(usedef, [], [e98])).
fof(e98, plain, (sP59 <=> (~ (e23 = h15(e14)) & ~ (e23 = h15(e13)) & ~ (e23 = h15(e12)) & ~ (e23 = h15(e11)) & ~ (e23 = h15(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP59])])).
fof(f99, plain, ((~ (e20 = h16(e14)) & ~ (e20 = h16(e13)) & ~ (e20 = h16(e12)) & ~ (e20 = h16(e11)) & ~ (e20 = h16(e10))) | ~ sP60), inference(usedef, [], [e99])).
fof(e99, plain, (sP60 <=> (~ (e20 = h16(e14)) & ~ (e20 = h16(e13)) & ~ (e20 = h16(e12)) & ~ (e20 = h16(e11)) & ~ (e20 = h16(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP60])])).
fof(f100, plain, ((~ (e21 = h16(e14)) & ~ (e21 = h16(e13)) & ~ (e21 = h16(e12)) & ~ (e21 = h16(e11)) & ~ (e21 = h16(e10))) | ~ sP61), inference(usedef, [], [e100])).
fof(e100, plain, (sP61 <=> (~ (e21 = h16(e14)) & ~ (e21 = h16(e13)) & ~ (e21 = h16(e12)) & ~ (e21 = h16(e11)) & ~ (e21 = h16(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP61])])).
fof(f101, plain, ((~ (e22 = h16(e14)) & ~ (e22 = h16(e13)) & ~ (e22 = h16(e12)) & ~ (e22 = h16(e11)) & ~ (e22 = h16(e10))) | ~ sP62), inference(usedef, [], [e101])).
fof(e101, plain, (sP62 <=> (~ (e22 = h16(e14)) & ~ (e22 = h16(e13)) & ~ (e22 = h16(e12)) & ~ (e22 = h16(e11)) & ~ (e22 = h16(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP62])])).
fof(f102, plain, ((~ (e23 = h16(e14)) & ~ (e23 = h16(e13)) & ~ (e23 = h16(e12)) & ~ (e23 = h16(e11)) & ~ (e23 = h16(e10))) | ~ sP63), inference(usedef, [], [e102])).
fof(e102, plain, (sP63 <=> (~ (e23 = h16(e14)) & ~ (e23 = h16(e13)) & ~ (e23 = h16(e12)) & ~ (e23 = h16(e11)) & ~ (e23 = h16(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP63])])).
fof(f103, plain, ((~ (e20 = h17(e14)) & ~ (e20 = h17(e13)) & ~ (e20 = h17(e12)) & ~ (e20 = h17(e11)) & ~ (e20 = h17(e10))) | ~ sP64), inference(usedef, [], [e103])).
fof(e103, plain, (sP64 <=> (~ (e20 = h17(e14)) & ~ (e20 = h17(e13)) & ~ (e20 = h17(e12)) & ~ (e20 = h17(e11)) & ~ (e20 = h17(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP64])])).
fof(f104, plain, ((~ (e21 = h17(e14)) & ~ (e21 = h17(e13)) & ~ (e21 = h17(e12)) & ~ (e21 = h17(e11)) & ~ (e21 = h17(e10))) | ~ sP65), inference(usedef, [], [e104])).
fof(e104, plain, (sP65 <=> (~ (e21 = h17(e14)) & ~ (e21 = h17(e13)) & ~ (e21 = h17(e12)) & ~ (e21 = h17(e11)) & ~ (e21 = h17(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP65])])).
fof(f105, plain, ((~ (e22 = h17(e14)) & ~ (e22 = h17(e13)) & ~ (e22 = h17(e12)) & ~ (e22 = h17(e11)) & ~ (e22 = h17(e10))) | ~ sP66), inference(usedef, [], [e105])).
fof(e105, plain, (sP66 <=> (~ (e22 = h17(e14)) & ~ (e22 = h17(e13)) & ~ (e22 = h17(e12)) & ~ (e22 = h17(e11)) & ~ (e22 = h17(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP66])])).
fof(f106, plain, ((~ (e23 = h17(e14)) & ~ (e23 = h17(e13)) & ~ (e23 = h17(e12)) & ~ (e23 = h17(e11)) & ~ (e23 = h17(e10))) | ~ sP67), inference(usedef, [], [e106])).
fof(e106, plain, (sP67 <=> (~ (e23 = h17(e14)) & ~ (e23 = h17(e13)) & ~ (e23 = h17(e12)) & ~ (e23 = h17(e11)) & ~ (e23 = h17(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP67])])).
fof(f107, plain, ((~ (e20 = h18(e14)) & ~ (e20 = h18(e13)) & ~ (e20 = h18(e12)) & ~ (e20 = h18(e11)) & ~ (e20 = h18(e10))) | ~ sP68), inference(usedef, [], [e107])).
fof(e107, plain, (sP68 <=> (~ (e20 = h18(e14)) & ~ (e20 = h18(e13)) & ~ (e20 = h18(e12)) & ~ (e20 = h18(e11)) & ~ (e20 = h18(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP68])])).
fof(f109, plain, ((~ (e22 = h18(e14)) & ~ (e22 = h18(e13)) & ~ (e22 = h18(e12)) & ~ (e22 = h18(e11)) & ~ (e22 = h18(e10))) | ~ sP70), inference(usedef, [], [e109])).
fof(e109, plain, (sP70 <=> (~ (e22 = h18(e14)) & ~ (e22 = h18(e13)) & ~ (e22 = h18(e12)) & ~ (e22 = h18(e11)) & ~ (e22 = h18(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP70])])).
fof(f110, plain, ((~ (e23 = h18(e14)) & ~ (e23 = h18(e13)) & ~ (e23 = h18(e12)) & ~ (e23 = h18(e11)) & ~ (e23 = h18(e10))) | ~ sP71), inference(usedef, [], [e110])).
fof(e110, plain, (sP71 <=> (~ (e23 = h18(e14)) & ~ (e23 = h18(e13)) & ~ (e23 = h18(e12)) & ~ (e23 = h18(e11)) & ~ (e23 = h18(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP71])])).
fof(f111, plain, ((~ (e20 = h19(e14)) & ~ (e20 = h19(e13)) & ~ (e20 = h19(e12)) & ~ (e20 = h19(e11)) & ~ (e20 = h19(e10))) | ~ sP72), inference(usedef, [], [e111])).
fof(e111, plain, (sP72 <=> (~ (e20 = h19(e14)) & ~ (e20 = h19(e13)) & ~ (e20 = h19(e12)) & ~ (e20 = h19(e11)) & ~ (e20 = h19(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP72])])).
fof(f112, plain, ((~ (e21 = h19(e14)) & ~ (e21 = h19(e13)) & ~ (e21 = h19(e12)) & ~ (e21 = h19(e11)) & ~ (e21 = h19(e10))) | ~ sP73), inference(usedef, [], [e112])).
fof(e112, plain, (sP73 <=> (~ (e21 = h19(e14)) & ~ (e21 = h19(e13)) & ~ (e21 = h19(e12)) & ~ (e21 = h19(e11)) & ~ (e21 = h19(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP73])])).
fof(f113, plain, ((~ (e22 = h19(e14)) & ~ (e22 = h19(e13)) & ~ (e22 = h19(e12)) & ~ (e22 = h19(e11)) & ~ (e22 = h19(e10))) | ~ sP74), inference(usedef, [], [e113])).
fof(e113, plain, (sP74 <=> (~ (e22 = h19(e14)) & ~ (e22 = h19(e13)) & ~ (e22 = h19(e12)) & ~ (e22 = h19(e11)) & ~ (e22 = h19(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP74])])).
fof(f114, plain, ((~ (e23 = h19(e14)) & ~ (e23 = h19(e13)) & ~ (e23 = h19(e12)) & ~ (e23 = h19(e11)) & ~ (e23 = h19(e10))) | ~ sP75), inference(usedef, [], [e114])).
fof(e114, plain, (sP75 <=> (~ (e23 = h19(e14)) & ~ (e23 = h19(e13)) & ~ (e23 = h19(e12)) & ~ (e23 = h19(e11)) & ~ (e23 = h19(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP75])])).
fof(f115, plain, ((~ (e20 = h20(e14)) & ~ (e20 = h20(e13)) & ~ (e20 = h20(e12)) & ~ (e20 = h20(e11)) & ~ (e20 = h20(e10))) | ~ sP76), inference(usedef, [], [e115])).
fof(e115, plain, (sP76 <=> (~ (e20 = h20(e14)) & ~ (e20 = h20(e13)) & ~ (e20 = h20(e12)) & ~ (e20 = h20(e11)) & ~ (e20 = h20(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP76])])).
fof(f116, plain, ((~ (e21 = h20(e14)) & ~ (e21 = h20(e13)) & ~ (e21 = h20(e12)) & ~ (e21 = h20(e11)) & ~ (e21 = h20(e10))) | ~ sP77), inference(usedef, [], [e116])).
fof(e116, plain, (sP77 <=> (~ (e21 = h20(e14)) & ~ (e21 = h20(e13)) & ~ (e21 = h20(e12)) & ~ (e21 = h20(e11)) & ~ (e21 = h20(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP77])])).
fof(f117, plain, ((~ (e22 = h20(e14)) & ~ (e22 = h20(e13)) & ~ (e22 = h20(e12)) & ~ (e22 = h20(e11)) & ~ (e22 = h20(e10))) | ~ sP78), inference(usedef, [], [e117])).
fof(e117, plain, (sP78 <=> (~ (e22 = h20(e14)) & ~ (e22 = h20(e13)) & ~ (e22 = h20(e12)) & ~ (e22 = h20(e11)) & ~ (e22 = h20(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP78])])).
fof(f118, plain, ((~ (e23 = h20(e14)) & ~ (e23 = h20(e13)) & ~ (e23 = h20(e12)) & ~ (e23 = h20(e11)) & ~ (e23 = h20(e10))) | ~ sP79), inference(usedef, [], [e118])).
fof(e118, plain, (sP79 <=> (~ (e23 = h20(e14)) & ~ (e23 = h20(e13)) & ~ (e23 = h20(e12)) & ~ (e23 = h20(e11)) & ~ (e23 = h20(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP79])])).
fof(f38, plain, (((~ (e24 = h20(e14)) & ~ (e24 = h20(e13)) & ~ (e24 = h20(e12)) & ~ (e24 = h20(e11)) & ~ (e24 = h20(e10))) | (~ (e23 = h20(e14)) & ~ (e23 = h20(e13)) & ~ (e23 = h20(e12)) & ~ (e23 = h20(e11)) & ~ (e23 = h20(e10))) | (~ (e22 = h20(e14)) & ~ (e22 = h20(e13)) & ~ (e22 = h20(e12)) & ~ (e22 = h20(e11)) & ~ (e22 = h20(e10))) | (~ (e21 = h20(e14)) & ~ (e21 = h20(e13)) & ~ (e21 = h20(e12)) & ~ (e21 = h20(e11)) & ~ (e21 = h20(e10))) | (~ (e20 = h20(e14)) & ~ (e20 = h20(e13)) & ~ (e20 = h20(e12)) & ~ (e20 = h20(e11)) & ~ (e20 = h20(e10))) | ~ (h20(op1(e14, e14)) = op2(h20(e14), h20(e14))) | ~ (h20(op1(e14, e13)) = op2(h20(e14), h20(e13))) | ~ (h20(op1(e14, e12)) = op2(h20(e14), h20(e12))) | ~ (h20(op1(e14, e11)) = op2(h20(e14), h20(e11))) | ~ (h20(op1(e14, e10)) = op2(h20(e14), h20(e10))) | ~ (h20(op1(e13, e14)) = op2(h20(e13), h20(e14))) | ~ (h20(op1(e13, e13)) = op2(h20(e13), h20(e13))) | ~ (h20(op1(e13, e12)) = op2(h20(e13), h20(e12))) | ~ (h20(op1(e13, e11)) = op2(h20(e13), h20(e11))) | ~ (h20(op1(e13, e10)) = op2(h20(e13), h20(e10))) | ~ (h20(op1(e12, e14)) = op2(h20(e12), h20(e14))) | ~ (h20(op1(e12, e13)) = op2(h20(e12), h20(e13))) | ~ (h20(op1(e12, e12)) = op2(h20(e12), h20(e12))) | ~ (h20(op1(e12, e11)) = op2(h20(e12), h20(e11))) | ~ (h20(op1(e12, e10)) = op2(h20(e12), h20(e10))) | ~ (h20(op1(e11, e14)) = op2(h20(e11), h20(e14))) | ~ (h20(op1(e11, e13)) = op2(h20(e11), h20(e13))) | ~ (h20(op1(e11, e12)) = op2(h20(e11), h20(e12))) | ~ (h20(op1(e11, e11)) = op2(h20(e11), h20(e11))) | ~ (h20(op1(e11, e10)) = op2(h20(e11), h20(e10))) | ~ (h20(op1(e10, e14)) = op2(h20(e10), h20(e14))) | ~ (h20(op1(e10, e13)) = op2(h20(e10), h20(e13))) | ~ (h20(op1(e10, e12)) = op2(h20(e10), h20(e12))) | ~ (h20(op1(e10, e11)) = op2(h20(e10), h20(e11))) | ~ (h20(op1(e10, e10)) = op2(h20(e10), h20(e10)))) & ((~ (e24 = h19(e14)) & ~ (e24 = h19(e13)) & ~ (e24 = h19(e12)) & ~ (e24 = h19(e11)) & ~ (e24 = h19(e10))) | (~ (e23 = h19(e14)) & ~ (e23 = h19(e13)) & ~ (e23 = h19(e12)) & ~ (e23 = h19(e11)) & ~ (e23 = h19(e10))) | (~ (e22 = h19(e14)) & ~ (e22 = h19(e13)) & ~ (e22 = h19(e12)) & ~ (e22 = h19(e11)) & ~ (e22 = h19(e10))) | (~ (e21 = h19(e14)) & ~ (e21 = h19(e13)) & ~ (e21 = h19(e12)) & ~ (e21 = h19(e11)) & ~ (e21 = h19(e10))) | (~ (e20 = h19(e14)) & ~ (e20 = h19(e13)) & ~ (e20 = h19(e12)) & ~ (e20 = h19(e11)) & ~ (e20 = h19(e10))) | ~ (h19(op1(e14, e14)) = op2(h19(e14), h19(e14))) | ~ (h19(op1(e14, e13)) = op2(h19(e14), h19(e13))) | ~ (h19(op1(e14, e12)) = op2(h19(e14), h19(e12))) | ~ (h19(op1(e14, e11)) = op2(h19(e14), h19(e11))) | ~ (h19(op1(e14, e10)) = op2(h19(e14), h19(e10))) | ~ (h19(op1(e13, e14)) = op2(h19(e13), h19(e14))) | ~ (h19(op1(e13, e13)) = op2(h19(e13), h19(e13))) | ~ (h19(op1(e13, e12)) = op2(h19(e13), h19(e12))) | ~ (h19(op1(e13, e11)) = op2(h19(e13), h19(e11))) | ~ (h19(op1(e13, e10)) = op2(h19(e13), h19(e10))) | ~ (h19(op1(e12, e14)) = op2(h19(e12), h19(e14))) | ~ (h19(op1(e12, e13)) = op2(h19(e12), h19(e13))) | ~ (h19(op1(e12, e12)) = op2(h19(e12), h19(e12))) | ~ (h19(op1(e12, e11)) = op2(h19(e12), h19(e11))) | ~ (h19(op1(e12, e10)) = op2(h19(e12), h19(e10))) | ~ (h19(op1(e11, e14)) = op2(h19(e11), h19(e14))) | ~ (h19(op1(e11, e13)) = op2(h19(e11), h19(e13))) | ~ (h19(op1(e11, e12)) = op2(h19(e11), h19(e12))) | ~ (h19(op1(e11, e11)) = op2(h19(e11), h19(e11))) | ~ (h19(op1(e11, e10)) = op2(h19(e11), h19(e10))) | ~ (h19(op1(e10, e14)) = op2(h19(e10), h19(e14))) | ~ (h19(op1(e10, e13)) = op2(h19(e10), h19(e13))) | ~ (h19(op1(e10, e12)) = op2(h19(e10), h19(e12))) | ~ (h19(op1(e10, e11)) = op2(h19(e10), h19(e11))) | ~ (h19(op1(e10, e10)) = op2(h19(e10), h19(e10)))) & ((~ (e24 = h18(e14)) & ~ (e24 = h18(e13)) & ~ (e24 = h18(e12)) & ~ (e24 = h18(e11)) & ~ (e24 = h18(e10))) | (~ (e23 = h18(e14)) & ~ (e23 = h18(e13)) & ~ (e23 = h18(e12)) & ~ (e23 = h18(e11)) & ~ (e23 = h18(e10))) | (~ (e22 = h18(e14)) & ~ (e22 = h18(e13)) & ~ (e22 = h18(e12)) & ~ (e22 = h18(e11)) & ~ (e22 = h18(e10))) | (~ (e21 = h18(e14)) & ~ (e21 = h18(e13)) & ~ (e21 = h18(e12)) & ~ (e21 = h18(e11)) & ~ (e21 = h18(e10))) | (~ (e20 = h18(e14)) & ~ (e20 = h18(e13)) & ~ (e20 = h18(e12)) & ~ (e20 = h18(e11)) & ~ (e20 = h18(e10))) | ~ (h18(op1(e14, e14)) = op2(h18(e14), h18(e14))) | ~ (h18(op1(e14, e13)) = op2(h18(e14), h18(e13))) | ~ (h18(op1(e14, e12)) = op2(h18(e14), h18(e12))) | ~ (h18(op1(e14, e11)) = op2(h18(e14), h18(e11))) | ~ (h18(op1(e14, e10)) = op2(h18(e14), h18(e10))) | ~ (h18(op1(e13, e14)) = op2(h18(e13), h18(e14))) | ~ (h18(op1(e13, e13)) = op2(h18(e13), h18(e13))) | ~ (h18(op1(e13, e12)) = op2(h18(e13), h18(e12))) | ~ (h18(op1(e13, e11)) = op2(h18(e13), h18(e11))) | ~ (h18(op1(e13, e10)) = op2(h18(e13), h18(e10))) | ~ (h18(op1(e12, e14)) = op2(h18(e12), h18(e14))) | ~ (h18(op1(e12, e13)) = op2(h18(e12), h18(e13))) | ~ (h18(op1(e12, e12)) = op2(h18(e12), h18(e12))) | ~ (h18(op1(e12, e11)) = op2(h18(e12), h18(e11))) | ~ (h18(op1(e12, e10)) = op2(h18(e12), h18(e10))) | ~ (h18(op1(e11, e14)) = op2(h18(e11), h18(e14))) | ~ (h18(op1(e11, e13)) = op2(h18(e11), h18(e13))) | ~ (h18(op1(e11, e12)) = op2(h18(e11), h18(e12))) | ~ (h18(op1(e11, e11)) = op2(h18(e11), h18(e11))) | ~ (h18(op1(e11, e10)) = op2(h18(e11), h18(e10))) | ~ (h18(op1(e10, e14)) = op2(h18(e10), h18(e14))) | ~ (h18(op1(e10, e13)) = op2(h18(e10), h18(e13))) | ~ (h18(op1(e10, e12)) = op2(h18(e10), h18(e12))) | ~ (h18(op1(e10, e11)) = op2(h18(e10), h18(e11))) | ~ (h18(op1(e10, e10)) = op2(h18(e10), h18(e10)))) & ((~ (e24 = h17(e14)) & ~ (e24 = h17(e13)) & ~ (e24 = h17(e12)) & ~ (e24 = h17(e11)) & ~ (e24 = h17(e10))) | (~ (e23 = h17(e14)) & ~ (e23 = h17(e13)) & ~ (e23 = h17(e12)) & ~ (e23 = h17(e11)) & ~ (e23 = h17(e10))) | (~ (e22 = h17(e14)) & ~ (e22 = h17(e13)) & ~ (e22 = h17(e12)) & ~ (e22 = h17(e11)) & ~ (e22 = h17(e10))) | (~ (e21 = h17(e14)) & ~ (e21 = h17(e13)) & ~ (e21 = h17(e12)) & ~ (e21 = h17(e11)) & ~ (e21 = h17(e10))) | (~ (e20 = h17(e14)) & ~ (e20 = h17(e13)) & ~ (e20 = h17(e12)) & ~ (e20 = h17(e11)) & ~ (e20 = h17(e10))) | ~ (h17(op1(e14, e14)) = op2(h17(e14), h17(e14))) | ~ (h17(op1(e14, e13)) = op2(h17(e14), h17(e13))) | ~ (h17(op1(e14, e12)) = op2(h17(e14), h17(e12))) | ~ (h17(op1(e14, e11)) = op2(h17(e14), h17(e11))) | ~ (h17(op1(e14, e10)) = op2(h17(e14), h17(e10))) | ~ (h17(op1(e13, e14)) = op2(h17(e13), h17(e14))) | ~ (h17(op1(e13, e13)) = op2(h17(e13), h17(e13))) | ~ (h17(op1(e13, e12)) = op2(h17(e13), h17(e12))) | ~ (h17(op1(e13, e11)) = op2(h17(e13), h17(e11))) | ~ (h17(op1(e13, e10)) = op2(h17(e13), h17(e10))) | ~ (h17(op1(e12, e14)) = op2(h17(e12), h17(e14))) | ~ (h17(op1(e12, e13)) = op2(h17(e12), h17(e13))) | ~ (h17(op1(e12, e12)) = op2(h17(e12), h17(e12))) | ~ (h17(op1(e12, e11)) = op2(h17(e12), h17(e11))) | ~ (h17(op1(e12, e10)) = op2(h17(e12), h17(e10))) | ~ (h17(op1(e11, e14)) = op2(h17(e11), h17(e14))) | ~ (h17(op1(e11, e13)) = op2(h17(e11), h17(e13))) | ~ (h17(op1(e11, e12)) = op2(h17(e11), h17(e12))) | ~ (h17(op1(e11, e11)) = op2(h17(e11), h17(e11))) | ~ (h17(op1(e11, e10)) = op2(h17(e11), h17(e10))) | ~ (h17(op1(e10, e14)) = op2(h17(e10), h17(e14))) | ~ (h17(op1(e10, e13)) = op2(h17(e10), h17(e13))) | ~ (h17(op1(e10, e12)) = op2(h17(e10), h17(e12))) | ~ (h17(op1(e10, e11)) = op2(h17(e10), h17(e11))) | ~ (h17(op1(e10, e10)) = op2(h17(e10), h17(e10)))) & ((~ (e24 = h16(e14)) & ~ (e24 = h16(e13)) & ~ (e24 = h16(e12)) & ~ (e24 = h16(e11)) & ~ (e24 = h16(e10))) | (~ (e23 = h16(e14)) & ~ (e23 = h16(e13)) & ~ (e23 = h16(e12)) & ~ (e23 = h16(e11)) & ~ (e23 = h16(e10))) | (~ (e22 = h16(e14)) & ~ (e22 = h16(e13)) & ~ (e22 = h16(e12)) & ~ (e22 = h16(e11)) & ~ (e22 = h16(e10))) | (~ (e21 = h16(e14)) & ~ (e21 = h16(e13)) & ~ (e21 = h16(e12)) & ~ (e21 = h16(e11)) & ~ (e21 = h16(e10))) | (~ (e20 = h16(e14)) & ~ (e20 = h16(e13)) & ~ (e20 = h16(e12)) & ~ (e20 = h16(e11)) & ~ (e20 = h16(e10))) | ~ (h16(op1(e14, e14)) = op2(h16(e14), h16(e14))) | ~ (h16(op1(e14, e13)) = op2(h16(e14), h16(e13))) | ~ (h16(op1(e14, e12)) = op2(h16(e14), h16(e12))) | ~ (h16(op1(e14, e11)) = op2(h16(e14), h16(e11))) | ~ (h16(op1(e14, e10)) = op2(h16(e14), h16(e10))) | ~ (h16(op1(e13, e14)) = op2(h16(e13), h16(e14))) | ~ (h16(op1(e13, e13)) = op2(h16(e13), h16(e13))) | ~ (h16(op1(e13, e12)) = op2(h16(e13), h16(e12))) | ~ (h16(op1(e13, e11)) = op2(h16(e13), h16(e11))) | ~ (h16(op1(e13, e10)) = op2(h16(e13), h16(e10))) | ~ (h16(op1(e12, e14)) = op2(h16(e12), h16(e14))) | ~ (h16(op1(e12, e13)) = op2(h16(e12), h16(e13))) | ~ (h16(op1(e12, e12)) = op2(h16(e12), h16(e12))) | ~ (h16(op1(e12, e11)) = op2(h16(e12), h16(e11))) | ~ (h16(op1(e12, e10)) = op2(h16(e12), h16(e10))) | ~ (h16(op1(e11, e14)) = op2(h16(e11), h16(e14))) | ~ (h16(op1(e11, e13)) = op2(h16(e11), h16(e13))) | ~ (h16(op1(e11, e12)) = op2(h16(e11), h16(e12))) | ~ (h16(op1(e11, e11)) = op2(h16(e11), h16(e11))) | ~ (h16(op1(e11, e10)) = op2(h16(e11), h16(e10))) | ~ (h16(op1(e10, e14)) = op2(h16(e10), h16(e14))) | ~ (h16(op1(e10, e13)) = op2(h16(e10), h16(e13))) | ~ (h16(op1(e10, e12)) = op2(h16(e10), h16(e12))) | ~ (h16(op1(e10, e11)) = op2(h16(e10), h16(e11))) | ~ (h16(op1(e10, e10)) = op2(h16(e10), h16(e10)))) & ((~ (e24 = h15(e14)) & ~ (e24 = h15(e13)) & ~ (e24 = h15(e12)) & ~ (e24 = h15(e11)) & ~ (e24 = h15(e10))) | (~ (e23 = h15(e14)) & ~ (e23 = h15(e13)) & ~ (e23 = h15(e12)) & ~ (e23 = h15(e11)) & ~ (e23 = h15(e10))) | (~ (e22 = h15(e14)) & ~ (e22 = h15(e13)) & ~ (e22 = h15(e12)) & ~ (e22 = h15(e11)) & ~ (e22 = h15(e10))) | (~ (e21 = h15(e14)) & ~ (e21 = h15(e13)) & ~ (e21 = h15(e12)) & ~ (e21 = h15(e11)) & ~ (e21 = h15(e10))) | (~ (e20 = h15(e14)) & ~ (e20 = h15(e13)) & ~ (e20 = h15(e12)) & ~ (e20 = h15(e11)) & ~ (e20 = h15(e10))) | ~ (h15(op1(e14, e14)) = op2(h15(e14), h15(e14))) | ~ (h15(op1(e14, e13)) = op2(h15(e14), h15(e13))) | ~ (h15(op1(e14, e12)) = op2(h15(e14), h15(e12))) | ~ (h15(op1(e14, e11)) = op2(h15(e14), h15(e11))) | ~ (h15(op1(e14, e10)) = op2(h15(e14), h15(e10))) | ~ (h15(op1(e13, e14)) = op2(h15(e13), h15(e14))) | ~ (h15(op1(e13, e13)) = op2(h15(e13), h15(e13))) | ~ (h15(op1(e13, e12)) = op2(h15(e13), h15(e12))) | ~ (h15(op1(e13, e11)) = op2(h15(e13), h15(e11))) | ~ (h15(op1(e13, e10)) = op2(h15(e13), h15(e10))) | ~ (h15(op1(e12, e14)) = op2(h15(e12), h15(e14))) | ~ (h15(op1(e12, e13)) = op2(h15(e12), h15(e13))) | ~ (h15(op1(e12, e12)) = op2(h15(e12), h15(e12))) | ~ (h15(op1(e12, e11)) = op2(h15(e12), h15(e11))) | ~ (h15(op1(e12, e10)) = op2(h15(e12), h15(e10))) | ~ (h15(op1(e11, e14)) = op2(h15(e11), h15(e14))) | ~ (h15(op1(e11, e13)) = op2(h15(e11), h15(e13))) | ~ (h15(op1(e11, e12)) = op2(h15(e11), h15(e12))) | ~ (h15(op1(e11, e11)) = op2(h15(e11), h15(e11))) | ~ (h15(op1(e11, e10)) = op2(h15(e11), h15(e10))) | ~ (h15(op1(e10, e14)) = op2(h15(e10), h15(e14))) | ~ (h15(op1(e10, e13)) = op2(h15(e10), h15(e13))) | ~ (h15(op1(e10, e12)) = op2(h15(e10), h15(e12))) | ~ (h15(op1(e10, e11)) = op2(h15(e10), h15(e11))) | ~ (h15(op1(e10, e10)) = op2(h15(e10), h15(e10)))) & ((~ (e24 = h14(e14)) & ~ (e24 = h14(e13)) & ~ (e24 = h14(e12)) & ~ (e24 = h14(e11)) & ~ (e24 = h14(e10))) | (~ (e23 = h14(e14)) & ~ (e23 = h14(e13)) & ~ (e23 = h14(e12)) & ~ (e23 = h14(e11)) & ~ (e23 = h14(e10))) | (~ (e22 = h14(e14)) & ~ (e22 = h14(e13)) & ~ (e22 = h14(e12)) & ~ (e22 = h14(e11)) & ~ (e22 = h14(e10))) | (~ (e21 = h14(e14)) & ~ (e21 = h14(e13)) & ~ (e21 = h14(e12)) & ~ (e21 = h14(e11)) & ~ (e21 = h14(e10))) | (~ (e20 = h14(e14)) & ~ (e20 = h14(e13)) & ~ (e20 = h14(e12)) & ~ (e20 = h14(e11)) & ~ (e20 = h14(e10))) | ~ (h14(op1(e14, e14)) = op2(h14(e14), h14(e14))) | ~ (h14(op1(e14, e13)) = op2(h14(e14), h14(e13))) | ~ (h14(op1(e14, e12)) = op2(h14(e14), h14(e12))) | ~ (h14(op1(e14, e11)) = op2(h14(e14), h14(e11))) | ~ (h14(op1(e14, e10)) = op2(h14(e14), h14(e10))) | ~ (h14(op1(e13, e14)) = op2(h14(e13), h14(e14))) | ~ (h14(op1(e13, e13)) = op2(h14(e13), h14(e13))) | ~ (h14(op1(e13, e12)) = op2(h14(e13), h14(e12))) | ~ (h14(op1(e13, e11)) = op2(h14(e13), h14(e11))) | ~ (h14(op1(e13, e10)) = op2(h14(e13), h14(e10))) | ~ (h14(op1(e12, e14)) = op2(h14(e12), h14(e14))) | ~ (h14(op1(e12, e13)) = op2(h14(e12), h14(e13))) | ~ (h14(op1(e12, e12)) = op2(h14(e12), h14(e12))) | ~ (h14(op1(e12, e11)) = op2(h14(e12), h14(e11))) | ~ (h14(op1(e12, e10)) = op2(h14(e12), h14(e10))) | ~ (h14(op1(e11, e14)) = op2(h14(e11), h14(e14))) | ~ (h14(op1(e11, e13)) = op2(h14(e11), h14(e13))) | ~ (h14(op1(e11, e12)) = op2(h14(e11), h14(e12))) | ~ (h14(op1(e11, e11)) = op2(h14(e11), h14(e11))) | ~ (h14(op1(e11, e10)) = op2(h14(e11), h14(e10))) | ~ (h14(op1(e10, e14)) = op2(h14(e10), h14(e14))) | ~ (h14(op1(e10, e13)) = op2(h14(e10), h14(e13))) | ~ (h14(op1(e10, e12)) = op2(h14(e10), h14(e12))) | ~ (h14(op1(e10, e11)) = op2(h14(e10), h14(e11))) | ~ (h14(op1(e10, e10)) = op2(h14(e10), h14(e10)))) & ((~ (e24 = h13(e14)) & ~ (e24 = h13(e13)) & ~ (e24 = h13(e12)) & ~ (e24 = h13(e11)) & ~ (e24 = h13(e10))) | (~ (e23 = h13(e14)) & ~ (e23 = h13(e13)) & ~ (e23 = h13(e12)) & ~ (e23 = h13(e11)) & ~ (e23 = h13(e10))) | (~ (e22 = h13(e14)) & ~ (e22 = h13(e13)) & ~ (e22 = h13(e12)) & ~ (e22 = h13(e11)) & ~ (e22 = h13(e10))) | (~ (e21 = h13(e14)) & ~ (e21 = h13(e13)) & ~ (e21 = h13(e12)) & ~ (e21 = h13(e11)) & ~ (e21 = h13(e10))) | (~ (e20 = h13(e14)) & ~ (e20 = h13(e13)) & ~ (e20 = h13(e12)) & ~ (e20 = h13(e11)) & ~ (e20 = h13(e10))) | ~ (h13(op1(e14, e14)) = op2(h13(e14), h13(e14))) | ~ (h13(op1(e14, e13)) = op2(h13(e14), h13(e13))) | ~ (h13(op1(e14, e12)) = op2(h13(e14), h13(e12))) | ~ (h13(op1(e14, e11)) = op2(h13(e14), h13(e11))) | ~ (h13(op1(e14, e10)) = op2(h13(e14), h13(e10))) | ~ (h13(op1(e13, e14)) = op2(h13(e13), h13(e14))) | ~ (h13(op1(e13, e13)) = op2(h13(e13), h13(e13))) | ~ (h13(op1(e13, e12)) = op2(h13(e13), h13(e12))) | ~ (h13(op1(e13, e11)) = op2(h13(e13), h13(e11))) | ~ (h13(op1(e13, e10)) = op2(h13(e13), h13(e10))) | ~ (h13(op1(e12, e14)) = op2(h13(e12), h13(e14))) | ~ (h13(op1(e12, e13)) = op2(h13(e12), h13(e13))) | ~ (h13(op1(e12, e12)) = op2(h13(e12), h13(e12))) | ~ (h13(op1(e12, e11)) = op2(h13(e12), h13(e11))) | ~ (h13(op1(e12, e10)) = op2(h13(e12), h13(e10))) | ~ (h13(op1(e11, e14)) = op2(h13(e11), h13(e14))) | ~ (h13(op1(e11, e13)) = op2(h13(e11), h13(e13))) | ~ (h13(op1(e11, e12)) = op2(h13(e11), h13(e12))) | ~ (h13(op1(e11, e11)) = op2(h13(e11), h13(e11))) | ~ (h13(op1(e11, e10)) = op2(h13(e11), h13(e10))) | ~ (h13(op1(e10, e14)) = op2(h13(e10), h13(e14))) | ~ (h13(op1(e10, e13)) = op2(h13(e10), h13(e13))) | ~ (h13(op1(e10, e12)) = op2(h13(e10), h13(e12))) | ~ (h13(op1(e10, e11)) = op2(h13(e10), h13(e11))) | ~ (h13(op1(e10, e10)) = op2(h13(e10), h13(e10)))) & ((~ (e24 = h12(e14)) & ~ (e24 = h12(e13)) & ~ (e24 = h12(e12)) & ~ (e24 = h12(e11)) & ~ (e24 = h12(e10))) | (~ (e23 = h12(e14)) & ~ (e23 = h12(e13)) & ~ (e23 = h12(e12)) & ~ (e23 = h12(e11)) & ~ (e23 = h12(e10))) | (~ (e22 = h12(e14)) & ~ (e22 = h12(e13)) & ~ (e22 = h12(e12)) & ~ (e22 = h12(e11)) & ~ (e22 = h12(e10))) | (~ (e21 = h12(e14)) & ~ (e21 = h12(e13)) & ~ (e21 = h12(e12)) & ~ (e21 = h12(e11)) & ~ (e21 = h12(e10))) | (~ (e20 = h12(e14)) & ~ (e20 = h12(e13)) & ~ (e20 = h12(e12)) & ~ (e20 = h12(e11)) & ~ (e20 = h12(e10))) | ~ (h12(op1(e14, e14)) = op2(h12(e14), h12(e14))) | ~ (h12(op1(e14, e13)) = op2(h12(e14), h12(e13))) | ~ (h12(op1(e14, e12)) = op2(h12(e14), h12(e12))) | ~ (h12(op1(e14, e11)) = op2(h12(e14), h12(e11))) | ~ (h12(op1(e14, e10)) = op2(h12(e14), h12(e10))) | ~ (h12(op1(e13, e14)) = op2(h12(e13), h12(e14))) | ~ (h12(op1(e13, e13)) = op2(h12(e13), h12(e13))) | ~ (h12(op1(e13, e12)) = op2(h12(e13), h12(e12))) | ~ (h12(op1(e13, e11)) = op2(h12(e13), h12(e11))) | ~ (h12(op1(e13, e10)) = op2(h12(e13), h12(e10))) | ~ (h12(op1(e12, e14)) = op2(h12(e12), h12(e14))) | ~ (h12(op1(e12, e13)) = op2(h12(e12), h12(e13))) | ~ (h12(op1(e12, e12)) = op2(h12(e12), h12(e12))) | ~ (h12(op1(e12, e11)) = op2(h12(e12), h12(e11))) | ~ (h12(op1(e12, e10)) = op2(h12(e12), h12(e10))) | ~ (h12(op1(e11, e14)) = op2(h12(e11), h12(e14))) | ~ (h12(op1(e11, e13)) = op2(h12(e11), h12(e13))) | ~ (h12(op1(e11, e12)) = op2(h12(e11), h12(e12))) | ~ (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) | ~ (h12(op1(e11, e10)) = op2(h12(e11), h12(e10))) | ~ (h12(op1(e10, e14)) = op2(h12(e10), h12(e14))) | ~ (h12(op1(e10, e13)) = op2(h12(e10), h12(e13))) | ~ (h12(op1(e10, e12)) = op2(h12(e10), h12(e12))) | ~ (h12(op1(e10, e11)) = op2(h12(e10), h12(e11))) | ~ (h12(op1(e10, e10)) = op2(h12(e10), h12(e10)))) & ((~ (e24 = h11(e14)) & ~ (e24 = h11(e13)) & ~ (e24 = h11(e12)) & ~ (e24 = h11(e11)) & ~ (e24 = h11(e10))) | (~ (e23 = h11(e14)) & ~ (e23 = h11(e13)) & ~ (e23 = h11(e12)) & ~ (e23 = h11(e11)) & ~ (e23 = h11(e10))) | (~ (e22 = h11(e14)) & ~ (e22 = h11(e13)) & ~ (e22 = h11(e12)) & ~ (e22 = h11(e11)) & ~ (e22 = h11(e10))) | (~ (e21 = h11(e14)) & ~ (e21 = h11(e13)) & ~ (e21 = h11(e12)) & ~ (e21 = h11(e11)) & ~ (e21 = h11(e10))) | (~ (e20 = h11(e14)) & ~ (e20 = h11(e13)) & ~ (e20 = h11(e12)) & ~ (e20 = h11(e11)) & ~ (e20 = h11(e10))) | ~ (h11(op1(e14, e14)) = op2(h11(e14), h11(e14))) | ~ (h11(op1(e14, e13)) = op2(h11(e14), h11(e13))) | ~ (h11(op1(e14, e12)) = op2(h11(e14), h11(e12))) | ~ (h11(op1(e14, e11)) = op2(h11(e14), h11(e11))) | ~ (h11(op1(e14, e10)) = op2(h11(e14), h11(e10))) | ~ (h11(op1(e13, e14)) = op2(h11(e13), h11(e14))) | ~ (h11(op1(e13, e13)) = op2(h11(e13), h11(e13))) | ~ (h11(op1(e13, e12)) = op2(h11(e13), h11(e12))) | ~ (h11(op1(e13, e11)) = op2(h11(e13), h11(e11))) | ~ (h11(op1(e13, e10)) = op2(h11(e13), h11(e10))) | ~ (h11(op1(e12, e14)) = op2(h11(e12), h11(e14))) | ~ (h11(op1(e12, e13)) = op2(h11(e12), h11(e13))) | ~ (h11(op1(e12, e12)) = op2(h11(e12), h11(e12))) | ~ (h11(op1(e12, e11)) = op2(h11(e12), h11(e11))) | ~ (h11(op1(e12, e10)) = op2(h11(e12), h11(e10))) | ~ (h11(op1(e11, e14)) = op2(h11(e11), h11(e14))) | ~ (h11(op1(e11, e13)) = op2(h11(e11), h11(e13))) | ~ (h11(op1(e11, e12)) = op2(h11(e11), h11(e12))) | ~ (h11(op1(e11, e11)) = op2(h11(e11), h11(e11))) | ~ (h11(op1(e11, e10)) = op2(h11(e11), h11(e10))) | ~ (h11(op1(e10, e14)) = op2(h11(e10), h11(e14))) | ~ (h11(op1(e10, e13)) = op2(h11(e10), h11(e13))) | ~ (h11(op1(e10, e12)) = op2(h11(e10), h11(e12))) | ~ (h11(op1(e10, e11)) = op2(h11(e10), h11(e11))) | ~ (h11(op1(e10, e10)) = op2(h11(e10), h11(e10)))) & ((~ (e24 = h10(e14)) & ~ (e24 = h10(e13)) & ~ (e24 = h10(e12)) & ~ (e24 = h10(e11)) & ~ (e24 = h10(e10))) | (~ (e23 = h10(e14)) & ~ (e23 = h10(e13)) & ~ (e23 = h10(e12)) & ~ (e23 = h10(e11)) & ~ (e23 = h10(e10))) | (~ (e22 = h10(e14)) & ~ (e22 = h10(e13)) & ~ (e22 = h10(e12)) & ~ (e22 = h10(e11)) & ~ (e22 = h10(e10))) | (~ (e21 = h10(e14)) & ~ (e21 = h10(e13)) & ~ (e21 = h10(e12)) & ~ (e21 = h10(e11)) & ~ (e21 = h10(e10))) | (~ (e20 = h10(e14)) & ~ (e20 = h10(e13)) & ~ (e20 = h10(e12)) & ~ (e20 = h10(e11)) & ~ (e20 = h10(e10))) | ~ (h10(op1(e14, e14)) = op2(h10(e14), h10(e14))) | ~ (h10(op1(e14, e13)) = op2(h10(e14), h10(e13))) | ~ (h10(op1(e14, e12)) = op2(h10(e14), h10(e12))) | ~ (h10(op1(e14, e11)) = op2(h10(e14), h10(e11))) | ~ (h10(op1(e14, e10)) = op2(h10(e14), h10(e10))) | ~ (h10(op1(e13, e14)) = op2(h10(e13), h10(e14))) | ~ (h10(op1(e13, e13)) = op2(h10(e13), h10(e13))) | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) | ~ (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), h10(e14))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) | ~ (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) | ~ (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) | ~ (h10(op1(e11, e14)) = op2(h10(e11), h10(e14))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))) & ((~ (e24 = h9(e14)) & ~ (e24 = h9(e13)) & ~ (e24 = h9(e12)) & ~ (e24 = h9(e11)) & ~ (e24 = h9(e10))) | (~ (e23 = h9(e14)) & ~ (e23 = h9(e13)) & ~ (e23 = h9(e12)) & ~ (e23 = h9(e11)) & ~ (e23 = h9(e10))) | (~ (e22 = h9(e14)) & ~ (e22 = h9(e13)) & ~ (e22 = h9(e12)) & ~ (e22 = h9(e11)) & ~ (e22 = h9(e10))) | (~ (e21 = h9(e14)) & ~ (e21 = h9(e13)) & ~ (e21 = h9(e12)) & ~ (e21 = h9(e11)) & ~ (e21 = h9(e10))) | (~ (e20 = h9(e14)) & ~ (e20 = h9(e13)) & ~ (e20 = h9(e12)) & ~ (e20 = h9(e11)) & ~ (e20 = h9(e10))) | ~ (h9(op1(e14, e14)) = op2(h9(e14), h9(e14))) | ~ (h9(op1(e14, e13)) = op2(h9(e14), h9(e13))) | ~ (h9(op1(e14, e12)) = op2(h9(e14), h9(e12))) | ~ (h9(op1(e14, e11)) = op2(h9(e14), h9(e11))) | ~ (h9(op1(e14, e10)) = op2(h9(e14), h9(e10))) | ~ (h9(op1(e13, e14)) = op2(h9(e13), h9(e14))) | ~ (h9(op1(e13, e13)) = op2(h9(e13), h9(e13))) | ~ (h9(op1(e13, e12)) = op2(h9(e13), h9(e12))) | ~ (h9(op1(e13, e11)) = op2(h9(e13), h9(e11))) | ~ (h9(op1(e13, e10)) = op2(h9(e13), h9(e10))) | ~ (h9(op1(e12, e14)) = op2(h9(e12), h9(e14))) | ~ (h9(op1(e12, e13)) = op2(h9(e12), h9(e13))) | ~ (h9(op1(e12, e12)) = op2(h9(e12), h9(e12))) | ~ (h9(op1(e12, e11)) = op2(h9(e12), h9(e11))) | ~ (h9(op1(e12, e10)) = op2(h9(e12), h9(e10))) | ~ (h9(op1(e11, e14)) = op2(h9(e11), h9(e14))) | ~ (h9(op1(e11, e13)) = op2(h9(e11), h9(e13))) | ~ (h9(op1(e11, e12)) = op2(h9(e11), h9(e12))) | ~ (h9(op1(e11, e11)) = op2(h9(e11), h9(e11))) | ~ (h9(op1(e11, e10)) = op2(h9(e11), h9(e10))) | ~ (h9(op1(e10, e14)) = op2(h9(e10), h9(e14))) | ~ (h9(op1(e10, e13)) = op2(h9(e10), h9(e13))) | ~ (h9(op1(e10, e12)) = op2(h9(e10), h9(e12))) | ~ (h9(op1(e10, e11)) = op2(h9(e10), h9(e11))) | ~ (h9(op1(e10, e10)) = op2(h9(e10), h9(e10)))) & ((~ (e24 = h8(e14)) & ~ (e24 = h8(e13)) & ~ (e24 = h8(e12)) & ~ (e24 = h8(e11)) & ~ (e24 = h8(e10))) | (~ (e23 = h8(e14)) & ~ (e23 = h8(e13)) & ~ (e23 = h8(e12)) & ~ (e23 = h8(e11)) & ~ (e23 = h8(e10))) | (~ (e22 = h8(e14)) & ~ (e22 = h8(e13)) & ~ (e22 = h8(e12)) & ~ (e22 = h8(e11)) & ~ (e22 = h8(e10))) | (~ (e21 = h8(e14)) & ~ (e21 = h8(e13)) & ~ (e21 = h8(e12)) & ~ (e21 = h8(e11)) & ~ (e21 = h8(e10))) | (~ (e20 = h8(e14)) & ~ (e20 = h8(e13)) & ~ (e20 = h8(e12)) & ~ (e20 = h8(e11)) & ~ (e20 = h8(e10))) | ~ (h8(op1(e14, e14)) = op2(h8(e14), h8(e14))) | ~ (h8(op1(e14, e13)) = op2(h8(e14), h8(e13))) | ~ (h8(op1(e14, e12)) = op2(h8(e14), h8(e12))) | ~ (h8(op1(e14, e11)) = op2(h8(e14), h8(e11))) | ~ (h8(op1(e14, e10)) = op2(h8(e14), h8(e10))) | ~ (h8(op1(e13, e14)) = op2(h8(e13), h8(e14))) | ~ (h8(op1(e13, e13)) = op2(h8(e13), h8(e13))) | ~ (h8(op1(e13, e12)) = op2(h8(e13), h8(e12))) | ~ (h8(op1(e13, e11)) = op2(h8(e13), h8(e11))) | ~ (h8(op1(e13, e10)) = op2(h8(e13), h8(e10))) | ~ (h8(op1(e12, e14)) = op2(h8(e12), h8(e14))) | ~ (h8(op1(e12, e13)) = op2(h8(e12), h8(e13))) | ~ (h8(op1(e12, e12)) = op2(h8(e12), h8(e12))) | ~ (h8(op1(e12, e11)) = op2(h8(e12), h8(e11))) | ~ (h8(op1(e12, e10)) = op2(h8(e12), h8(e10))) | ~ (h8(op1(e11, e14)) = op2(h8(e11), h8(e14))) | ~ (h8(op1(e11, e13)) = op2(h8(e11), h8(e13))) | ~ (h8(op1(e11, e12)) = op2(h8(e11), h8(e12))) | ~ (h8(op1(e11, e11)) = op2(h8(e11), h8(e11))) | ~ (h8(op1(e11, e10)) = op2(h8(e11), h8(e10))) | ~ (h8(op1(e10, e14)) = op2(h8(e10), h8(e14))) | ~ (h8(op1(e10, e13)) = op2(h8(e10), h8(e13))) | ~ (h8(op1(e10, e12)) = op2(h8(e10), h8(e12))) | ~ (h8(op1(e10, e11)) = op2(h8(e10), h8(e11))) | ~ (h8(op1(e10, e10)) = op2(h8(e10), h8(e10)))) & ((~ (e24 = h7(e14)) & ~ (e24 = h7(e13)) & ~ (e24 = h7(e12)) & ~ (e24 = h7(e11)) & ~ (e24 = h7(e10))) | (~ (e23 = h7(e14)) & ~ (e23 = h7(e13)) & ~ (e23 = h7(e12)) & ~ (e23 = h7(e11)) & ~ (e23 = h7(e10))) | (~ (e22 = h7(e14)) & ~ (e22 = h7(e13)) & ~ (e22 = h7(e12)) & ~ (e22 = h7(e11)) & ~ (e22 = h7(e10))) | (~ (e21 = h7(e14)) & ~ (e21 = h7(e13)) & ~ (e21 = h7(e12)) & ~ (e21 = h7(e11)) & ~ (e21 = h7(e10))) | (~ (e20 = h7(e14)) & ~ (e20 = h7(e13)) & ~ (e20 = h7(e12)) & ~ (e20 = h7(e11)) & ~ (e20 = h7(e10))) | ~ (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) | ~ (h7(op1(e14, e13)) = op2(h7(e14), h7(e13))) | ~ (h7(op1(e14, e12)) = op2(h7(e14), h7(e12))) | ~ (h7(op1(e14, e11)) = op2(h7(e14), h7(e11))) | ~ (h7(op1(e14, e10)) = op2(h7(e14), h7(e10))) | ~ (h7(op1(e13, e14)) = op2(h7(e13), h7(e14))) | ~ (h7(op1(e13, e13)) = op2(h7(e13), h7(e13))) | ~ (h7(op1(e13, e12)) = op2(h7(e13), h7(e12))) | ~ (h7(op1(e13, e11)) = op2(h7(e13), h7(e11))) | ~ (h7(op1(e13, e10)) = op2(h7(e13), h7(e10))) | ~ (h7(op1(e12, e14)) = op2(h7(e12), h7(e14))) | ~ (h7(op1(e12, e13)) = op2(h7(e12), h7(e13))) | ~ (h7(op1(e12, e12)) = op2(h7(e12), h7(e12))) | ~ (h7(op1(e12, e11)) = op2(h7(e12), h7(e11))) | ~ (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) | ~ (h7(op1(e11, e14)) = op2(h7(e11), h7(e14))) | ~ (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) | ~ (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) | ~ (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) | ~ (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) | ~ (h7(op1(e10, e14)) = op2(h7(e10), h7(e14))) | ~ (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) | ~ (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) | ~ (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))) & ((~ (e24 = h6(e14)) & ~ (e24 = h6(e13)) & ~ (e24 = h6(e12)) & ~ (e24 = h6(e11)) & ~ (e24 = h6(e10))) | (~ (e23 = h6(e14)) & ~ (e23 = h6(e13)) & ~ (e23 = h6(e12)) & ~ (e23 = h6(e11)) & ~ (e23 = h6(e10))) | (~ (e22 = h6(e14)) & ~ (e22 = h6(e13)) & ~ (e22 = h6(e12)) & ~ (e22 = h6(e11)) & ~ (e22 = h6(e10))) | (~ (e21 = h6(e14)) & ~ (e21 = h6(e13)) & ~ (e21 = h6(e12)) & ~ (e21 = h6(e11)) & ~ (e21 = h6(e10))) | (~ (e20 = h6(e14)) & ~ (e20 = h6(e13)) & ~ (e20 = h6(e12)) & ~ (e20 = h6(e11)) & ~ (e20 = h6(e10))) | ~ (h6(op1(e14, e14)) = op2(h6(e14), h6(e14))) | ~ (h6(op1(e14, e13)) = op2(h6(e14), h6(e13))) | ~ (h6(op1(e14, e12)) = op2(h6(e14), h6(e12))) | ~ (h6(op1(e14, e11)) = op2(h6(e14), h6(e11))) | ~ (h6(op1(e14, e10)) = op2(h6(e14), h6(e10))) | ~ (h6(op1(e13, e14)) = op2(h6(e13), h6(e14))) | ~ (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) | ~ (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))) & ((~ (e24 = h5(e14)) & ~ (e24 = h5(e13)) & ~ (e24 = h5(e12)) & ~ (e24 = h5(e11)) & ~ (e24 = h5(e10))) | (~ (e23 = h5(e14)) & ~ (e23 = h5(e13)) & ~ (e23 = h5(e12)) & ~ (e23 = h5(e11)) & ~ (e23 = h5(e10))) | (~ (e22 = h5(e14)) & ~ (e22 = h5(e13)) & ~ (e22 = h5(e12)) & ~ (e22 = h5(e11)) & ~ (e22 = h5(e10))) | (~ (e21 = h5(e14)) & ~ (e21 = h5(e13)) & ~ (e21 = h5(e12)) & ~ (e21 = h5(e11)) & ~ (e21 = h5(e10))) | (~ (e20 = h5(e14)) & ~ (e20 = h5(e13)) & ~ (e20 = h5(e12)) & ~ (e20 = h5(e11)) & ~ (e20 = h5(e10))) | ~ (h5(op1(e14, e14)) = op2(h5(e14), h5(e14))) | ~ (h5(op1(e14, e13)) = op2(h5(e14), h5(e13))) | ~ (h5(op1(e14, e12)) = op2(h5(e14), h5(e12))) | ~ (h5(op1(e14, e11)) = op2(h5(e14), h5(e11))) | ~ (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) | ~ (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))) & ((~ (e24 = h4(e14)) & ~ (e24 = h4(e13)) & ~ (e24 = h4(e12)) & ~ (e24 = h4(e11)) & ~ (e24 = h4(e10))) | (~ (e23 = h4(e14)) & ~ (e23 = h4(e13)) & ~ (e23 = h4(e12)) & ~ (e23 = h4(e11)) & ~ (e23 = h4(e10))) | (~ (e22 = h4(e14)) & ~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10))) | (~ (e21 = h4(e14)) & ~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10))) | (~ (e20 = h4(e14)) & ~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10))) | ~ (h4(op1(e14, e14)) = op2(h4(e14), h4(e14))) | ~ (h4(op1(e14, e13)) = op2(h4(e14), h4(e13))) | ~ (h4(op1(e14, e12)) = op2(h4(e14), h4(e12))) | ~ (h4(op1(e14, e11)) = op2(h4(e14), h4(e11))) | ~ (h4(op1(e14, e10)) = op2(h4(e14), h4(e10))) | ~ (h4(op1(e13, e14)) = op2(h4(e13), h4(e14))) | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e14)) = op2(h4(e12), h4(e14))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e14)) = op2(h4(e11), h4(e14))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e14)) = op2(h4(e10), h4(e14))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) & ((~ (e24 = h3(e14)) & ~ (e24 = h3(e13)) & ~ (e24 = h3(e12)) & ~ (e24 = h3(e11)) & ~ (e24 = h3(e10))) | (~ (e23 = h3(e14)) & ~ (e23 = h3(e13)) & ~ (e23 = h3(e12)) & ~ (e23 = h3(e11)) & ~ (e23 = h3(e10))) | (~ (e22 = h3(e14)) & ~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10))) | (~ (e21 = h3(e14)) & ~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10))) | (~ (e20 = h3(e14)) & ~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10))) | ~ (h3(op1(e14, e14)) = op2(h3(e14), h3(e14))) | ~ (h3(op1(e14, e13)) = op2(h3(e14), h3(e13))) | ~ (h3(op1(e14, e12)) = op2(h3(e14), h3(e12))) | ~ (h3(op1(e14, e11)) = op2(h3(e14), h3(e11))) | ~ (h3(op1(e14, e10)) = op2(h3(e14), h3(e10))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), h3(e14))) | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) & ((~ (e24 = h2(e14)) & ~ (e24 = h2(e13)) & ~ (e24 = h2(e12)) & ~ (e24 = h2(e11)) & ~ (e24 = h2(e10))) | (~ (e23 = h2(e14)) & ~ (e23 = h2(e13)) & ~ (e23 = h2(e12)) & ~ (e23 = h2(e11)) & ~ (e23 = h2(e10))) | (~ (e22 = h2(e14)) & ~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10))) | (~ (e21 = h2(e14)) & ~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10))) | (~ (e20 = h2(e14)) & ~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10))) | ~ (h2(op1(e14, e14)) = op2(h2(e14), h2(e14))) | ~ (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) | ~ (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) | ~ (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) | ~ (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) | ~ (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) & ((~ (e24 = h1(e14)) & ~ (e24 = h1(e13)) & ~ (e24 = h1(e12)) & ~ (e24 = h1(e11)) & ~ (e24 = h1(e10))) | (~ (e23 = h1(e14)) & ~ (e23 = h1(e13)) & ~ (e23 = h1(e12)) & ~ (e23 = h1(e11)) & ~ (e23 = h1(e10))) | (~ (e22 = h1(e14)) & ~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10))) | (~ (e21 = h1(e14)) & ~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10))) | (~ (e20 = h1(e14)) & ~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10))) | ~ (h1(op1(e14, e14)) = op2(h1(e14), h1(e14))) | ~ (h1(op1(e14, e13)) = op2(h1(e14), h1(e13))) | ~ (h1(op1(e14, e12)) = op2(h1(e14), h1(e12))) | ~ (h1(op1(e14, e11)) = op2(h1(e14), h1(e11))) | ~ (h1(op1(e14, e10)) = op2(h1(e14), h1(e10))) | ~ (h1(op1(e13, e14)) = op2(h1(e13), h1(e14))) | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e14)) = op2(h1(e12), h1(e14))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e14)) = op2(h1(e11), h1(e14))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e14)) = op2(h1(e10), h1(e14))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(ennf_transformation, [], [f37])).
fof(f37, plain, ~ ((((e24 = h20(e14)) | (e24 = h20(e13)) | (e24 = h20(e12)) | (e24 = h20(e11)) | (e24 = h20(e10))) & ((e23 = h20(e14)) | (e23 = h20(e13)) | (e23 = h20(e12)) | (e23 = h20(e11)) | (e23 = h20(e10))) & ((e22 = h20(e14)) | (e22 = h20(e13)) | (e22 = h20(e12)) | (e22 = h20(e11)) | (e22 = h20(e10))) & ((e21 = h20(e14)) | (e21 = h20(e13)) | (e21 = h20(e12)) | (e21 = h20(e11)) | (e21 = h20(e10))) & ((e20 = h20(e14)) | (e20 = h20(e13)) | (e20 = h20(e12)) | (e20 = h20(e11)) | (e20 = h20(e10))) & (h20(op1(e14, e14)) = op2(h20(e14), h20(e14))) & (h20(op1(e14, e13)) = op2(h20(e14), h20(e13))) & (h20(op1(e14, e12)) = op2(h20(e14), h20(e12))) & (h20(op1(e14, e11)) = op2(h20(e14), h20(e11))) & (h20(op1(e14, e10)) = op2(h20(e14), h20(e10))) & (h20(op1(e13, e14)) = op2(h20(e13), h20(e14))) & (h20(op1(e13, e13)) = op2(h20(e13), h20(e13))) & (h20(op1(e13, e12)) = op2(h20(e13), h20(e12))) & (h20(op1(e13, e11)) = op2(h20(e13), h20(e11))) & (h20(op1(e13, e10)) = op2(h20(e13), h20(e10))) & (h20(op1(e12, e14)) = op2(h20(e12), h20(e14))) & (h20(op1(e12, e13)) = op2(h20(e12), h20(e13))) & (h20(op1(e12, e12)) = op2(h20(e12), h20(e12))) & (h20(op1(e12, e11)) = op2(h20(e12), h20(e11))) & (h20(op1(e12, e10)) = op2(h20(e12), h20(e10))) & (h20(op1(e11, e14)) = op2(h20(e11), h20(e14))) & (h20(op1(e11, e13)) = op2(h20(e11), h20(e13))) & (h20(op1(e11, e12)) = op2(h20(e11), h20(e12))) & (h20(op1(e11, e11)) = op2(h20(e11), h20(e11))) & (h20(op1(e11, e10)) = op2(h20(e11), h20(e10))) & (h20(op1(e10, e14)) = op2(h20(e10), h20(e14))) & (h20(op1(e10, e13)) = op2(h20(e10), h20(e13))) & (h20(op1(e10, e12)) = op2(h20(e10), h20(e12))) & (h20(op1(e10, e11)) = op2(h20(e10), h20(e11))) & (h20(op1(e10, e10)) = op2(h20(e10), h20(e10)))) | (((e24 = h19(e14)) | (e24 = h19(e13)) | (e24 = h19(e12)) | (e24 = h19(e11)) | (e24 = h19(e10))) & ((e23 = h19(e14)) | (e23 = h19(e13)) | (e23 = h19(e12)) | (e23 = h19(e11)) | (e23 = h19(e10))) & ((e22 = h19(e14)) | (e22 = h19(e13)) | (e22 = h19(e12)) | (e22 = h19(e11)) | (e22 = h19(e10))) & ((e21 = h19(e14)) | (e21 = h19(e13)) | (e21 = h19(e12)) | (e21 = h19(e11)) | (e21 = h19(e10))) & ((e20 = h19(e14)) | (e20 = h19(e13)) | (e20 = h19(e12)) | (e20 = h19(e11)) | (e20 = h19(e10))) & (h19(op1(e14, e14)) = op2(h19(e14), h19(e14))) & (h19(op1(e14, e13)) = op2(h19(e14), h19(e13))) & (h19(op1(e14, e12)) = op2(h19(e14), h19(e12))) & (h19(op1(e14, e11)) = op2(h19(e14), h19(e11))) & (h19(op1(e14, e10)) = op2(h19(e14), h19(e10))) & (h19(op1(e13, e14)) = op2(h19(e13), h19(e14))) & (h19(op1(e13, e13)) = op2(h19(e13), h19(e13))) & (h19(op1(e13, e12)) = op2(h19(e13), h19(e12))) & (h19(op1(e13, e11)) = op2(h19(e13), h19(e11))) & (h19(op1(e13, e10)) = op2(h19(e13), h19(e10))) & (h19(op1(e12, e14)) = op2(h19(e12), h19(e14))) & (h19(op1(e12, e13)) = op2(h19(e12), h19(e13))) & (h19(op1(e12, e12)) = op2(h19(e12), h19(e12))) & (h19(op1(e12, e11)) = op2(h19(e12), h19(e11))) & (h19(op1(e12, e10)) = op2(h19(e12), h19(e10))) & (h19(op1(e11, e14)) = op2(h19(e11), h19(e14))) & (h19(op1(e11, e13)) = op2(h19(e11), h19(e13))) & (h19(op1(e11, e12)) = op2(h19(e11), h19(e12))) & (h19(op1(e11, e11)) = op2(h19(e11), h19(e11))) & (h19(op1(e11, e10)) = op2(h19(e11), h19(e10))) & (h19(op1(e10, e14)) = op2(h19(e10), h19(e14))) & (h19(op1(e10, e13)) = op2(h19(e10), h19(e13))) & (h19(op1(e10, e12)) = op2(h19(e10), h19(e12))) & (h19(op1(e10, e11)) = op2(h19(e10), h19(e11))) & (h19(op1(e10, e10)) = op2(h19(e10), h19(e10)))) | (((e24 = h18(e14)) | (e24 = h18(e13)) | (e24 = h18(e12)) | (e24 = h18(e11)) | (e24 = h18(e10))) & ((e23 = h18(e14)) | (e23 = h18(e13)) | (e23 = h18(e12)) | (e23 = h18(e11)) | (e23 = h18(e10))) & ((e22 = h18(e14)) | (e22 = h18(e13)) | (e22 = h18(e12)) | (e22 = h18(e11)) | (e22 = h18(e10))) & ((e21 = h18(e14)) | (e21 = h18(e13)) | (e21 = h18(e12)) | (e21 = h18(e11)) | (e21 = h18(e10))) & ((e20 = h18(e14)) | (e20 = h18(e13)) | (e20 = h18(e12)) | (e20 = h18(e11)) | (e20 = h18(e10))) & (h18(op1(e14, e14)) = op2(h18(e14), h18(e14))) & (h18(op1(e14, e13)) = op2(h18(e14), h18(e13))) & (h18(op1(e14, e12)) = op2(h18(e14), h18(e12))) & (h18(op1(e14, e11)) = op2(h18(e14), h18(e11))) & (h18(op1(e14, e10)) = op2(h18(e14), h18(e10))) & (h18(op1(e13, e14)) = op2(h18(e13), h18(e14))) & (h18(op1(e13, e13)) = op2(h18(e13), h18(e13))) & (h18(op1(e13, e12)) = op2(h18(e13), h18(e12))) & (h18(op1(e13, e11)) = op2(h18(e13), h18(e11))) & (h18(op1(e13, e10)) = op2(h18(e13), h18(e10))) & (h18(op1(e12, e14)) = op2(h18(e12), h18(e14))) & (h18(op1(e12, e13)) = op2(h18(e12), h18(e13))) & (h18(op1(e12, e12)) = op2(h18(e12), h18(e12))) & (h18(op1(e12, e11)) = op2(h18(e12), h18(e11))) & (h18(op1(e12, e10)) = op2(h18(e12), h18(e10))) & (h18(op1(e11, e14)) = op2(h18(e11), h18(e14))) & (h18(op1(e11, e13)) = op2(h18(e11), h18(e13))) & (h18(op1(e11, e12)) = op2(h18(e11), h18(e12))) & (h18(op1(e11, e11)) = op2(h18(e11), h18(e11))) & (h18(op1(e11, e10)) = op2(h18(e11), h18(e10))) & (h18(op1(e10, e14)) = op2(h18(e10), h18(e14))) & (h18(op1(e10, e13)) = op2(h18(e10), h18(e13))) & (h18(op1(e10, e12)) = op2(h18(e10), h18(e12))) & (h18(op1(e10, e11)) = op2(h18(e10), h18(e11))) & (h18(op1(e10, e10)) = op2(h18(e10), h18(e10)))) | (((e24 = h17(e14)) | (e24 = h17(e13)) | (e24 = h17(e12)) | (e24 = h17(e11)) | (e24 = h17(e10))) & ((e23 = h17(e14)) | (e23 = h17(e13)) | (e23 = h17(e12)) | (e23 = h17(e11)) | (e23 = h17(e10))) & ((e22 = h17(e14)) | (e22 = h17(e13)) | (e22 = h17(e12)) | (e22 = h17(e11)) | (e22 = h17(e10))) & ((e21 = h17(e14)) | (e21 = h17(e13)) | (e21 = h17(e12)) | (e21 = h17(e11)) | (e21 = h17(e10))) & ((e20 = h17(e14)) | (e20 = h17(e13)) | (e20 = h17(e12)) | (e20 = h17(e11)) | (e20 = h17(e10))) & (h17(op1(e14, e14)) = op2(h17(e14), h17(e14))) & (h17(op1(e14, e13)) = op2(h17(e14), h17(e13))) & (h17(op1(e14, e12)) = op2(h17(e14), h17(e12))) & (h17(op1(e14, e11)) = op2(h17(e14), h17(e11))) & (h17(op1(e14, e10)) = op2(h17(e14), h17(e10))) & (h17(op1(e13, e14)) = op2(h17(e13), h17(e14))) & (h17(op1(e13, e13)) = op2(h17(e13), h17(e13))) & (h17(op1(e13, e12)) = op2(h17(e13), h17(e12))) & (h17(op1(e13, e11)) = op2(h17(e13), h17(e11))) & (h17(op1(e13, e10)) = op2(h17(e13), h17(e10))) & (h17(op1(e12, e14)) = op2(h17(e12), h17(e14))) & (h17(op1(e12, e13)) = op2(h17(e12), h17(e13))) & (h17(op1(e12, e12)) = op2(h17(e12), h17(e12))) & (h17(op1(e12, e11)) = op2(h17(e12), h17(e11))) & (h17(op1(e12, e10)) = op2(h17(e12), h17(e10))) & (h17(op1(e11, e14)) = op2(h17(e11), h17(e14))) & (h17(op1(e11, e13)) = op2(h17(e11), h17(e13))) & (h17(op1(e11, e12)) = op2(h17(e11), h17(e12))) & (h17(op1(e11, e11)) = op2(h17(e11), h17(e11))) & (h17(op1(e11, e10)) = op2(h17(e11), h17(e10))) & (h17(op1(e10, e14)) = op2(h17(e10), h17(e14))) & (h17(op1(e10, e13)) = op2(h17(e10), h17(e13))) & (h17(op1(e10, e12)) = op2(h17(e10), h17(e12))) & (h17(op1(e10, e11)) = op2(h17(e10), h17(e11))) & (h17(op1(e10, e10)) = op2(h17(e10), h17(e10)))) | (((e24 = h16(e14)) | (e24 = h16(e13)) | (e24 = h16(e12)) | (e24 = h16(e11)) | (e24 = h16(e10))) & ((e23 = h16(e14)) | (e23 = h16(e13)) | (e23 = h16(e12)) | (e23 = h16(e11)) | (e23 = h16(e10))) & ((e22 = h16(e14)) | (e22 = h16(e13)) | (e22 = h16(e12)) | (e22 = h16(e11)) | (e22 = h16(e10))) & ((e21 = h16(e14)) | (e21 = h16(e13)) | (e21 = h16(e12)) | (e21 = h16(e11)) | (e21 = h16(e10))) & ((e20 = h16(e14)) | (e20 = h16(e13)) | (e20 = h16(e12)) | (e20 = h16(e11)) | (e20 = h16(e10))) & (h16(op1(e14, e14)) = op2(h16(e14), h16(e14))) & (h16(op1(e14, e13)) = op2(h16(e14), h16(e13))) & (h16(op1(e14, e12)) = op2(h16(e14), h16(e12))) & (h16(op1(e14, e11)) = op2(h16(e14), h16(e11))) & (h16(op1(e14, e10)) = op2(h16(e14), h16(e10))) & (h16(op1(e13, e14)) = op2(h16(e13), h16(e14))) & (h16(op1(e13, e13)) = op2(h16(e13), h16(e13))) & (h16(op1(e13, e12)) = op2(h16(e13), h16(e12))) & (h16(op1(e13, e11)) = op2(h16(e13), h16(e11))) & (h16(op1(e13, e10)) = op2(h16(e13), h16(e10))) & (h16(op1(e12, e14)) = op2(h16(e12), h16(e14))) & (h16(op1(e12, e13)) = op2(h16(e12), h16(e13))) & (h16(op1(e12, e12)) = op2(h16(e12), h16(e12))) & (h16(op1(e12, e11)) = op2(h16(e12), h16(e11))) & (h16(op1(e12, e10)) = op2(h16(e12), h16(e10))) & (h16(op1(e11, e14)) = op2(h16(e11), h16(e14))) & (h16(op1(e11, e13)) = op2(h16(e11), h16(e13))) & (h16(op1(e11, e12)) = op2(h16(e11), h16(e12))) & (h16(op1(e11, e11)) = op2(h16(e11), h16(e11))) & (h16(op1(e11, e10)) = op2(h16(e11), h16(e10))) & (h16(op1(e10, e14)) = op2(h16(e10), h16(e14))) & (h16(op1(e10, e13)) = op2(h16(e10), h16(e13))) & (h16(op1(e10, e12)) = op2(h16(e10), h16(e12))) & (h16(op1(e10, e11)) = op2(h16(e10), h16(e11))) & (h16(op1(e10, e10)) = op2(h16(e10), h16(e10)))) | (((e24 = h15(e14)) | (e24 = h15(e13)) | (e24 = h15(e12)) | (e24 = h15(e11)) | (e24 = h15(e10))) & ((e23 = h15(e14)) | (e23 = h15(e13)) | (e23 = h15(e12)) | (e23 = h15(e11)) | (e23 = h15(e10))) & ((e22 = h15(e14)) | (e22 = h15(e13)) | (e22 = h15(e12)) | (e22 = h15(e11)) | (e22 = h15(e10))) & ((e21 = h15(e14)) | (e21 = h15(e13)) | (e21 = h15(e12)) | (e21 = h15(e11)) | (e21 = h15(e10))) & ((e20 = h15(e14)) | (e20 = h15(e13)) | (e20 = h15(e12)) | (e20 = h15(e11)) | (e20 = h15(e10))) & (h15(op1(e14, e14)) = op2(h15(e14), h15(e14))) & (h15(op1(e14, e13)) = op2(h15(e14), h15(e13))) & (h15(op1(e14, e12)) = op2(h15(e14), h15(e12))) & (h15(op1(e14, e11)) = op2(h15(e14), h15(e11))) & (h15(op1(e14, e10)) = op2(h15(e14), h15(e10))) & (h15(op1(e13, e14)) = op2(h15(e13), h15(e14))) & (h15(op1(e13, e13)) = op2(h15(e13), h15(e13))) & (h15(op1(e13, e12)) = op2(h15(e13), h15(e12))) & (h15(op1(e13, e11)) = op2(h15(e13), h15(e11))) & (h15(op1(e13, e10)) = op2(h15(e13), h15(e10))) & (h15(op1(e12, e14)) = op2(h15(e12), h15(e14))) & (h15(op1(e12, e13)) = op2(h15(e12), h15(e13))) & (h15(op1(e12, e12)) = op2(h15(e12), h15(e12))) & (h15(op1(e12, e11)) = op2(h15(e12), h15(e11))) & (h15(op1(e12, e10)) = op2(h15(e12), h15(e10))) & (h15(op1(e11, e14)) = op2(h15(e11), h15(e14))) & (h15(op1(e11, e13)) = op2(h15(e11), h15(e13))) & (h15(op1(e11, e12)) = op2(h15(e11), h15(e12))) & (h15(op1(e11, e11)) = op2(h15(e11), h15(e11))) & (h15(op1(e11, e10)) = op2(h15(e11), h15(e10))) & (h15(op1(e10, e14)) = op2(h15(e10), h15(e14))) & (h15(op1(e10, e13)) = op2(h15(e10), h15(e13))) & (h15(op1(e10, e12)) = op2(h15(e10), h15(e12))) & (h15(op1(e10, e11)) = op2(h15(e10), h15(e11))) & (h15(op1(e10, e10)) = op2(h15(e10), h15(e10)))) | (((e24 = h14(e14)) | (e24 = h14(e13)) | (e24 = h14(e12)) | (e24 = h14(e11)) | (e24 = h14(e10))) & ((e23 = h14(e14)) | (e23 = h14(e13)) | (e23 = h14(e12)) | (e23 = h14(e11)) | (e23 = h14(e10))) & ((e22 = h14(e14)) | (e22 = h14(e13)) | (e22 = h14(e12)) | (e22 = h14(e11)) | (e22 = h14(e10))) & ((e21 = h14(e14)) | (e21 = h14(e13)) | (e21 = h14(e12)) | (e21 = h14(e11)) | (e21 = h14(e10))) & ((e20 = h14(e14)) | (e20 = h14(e13)) | (e20 = h14(e12)) | (e20 = h14(e11)) | (e20 = h14(e10))) & (h14(op1(e14, e14)) = op2(h14(e14), h14(e14))) & (h14(op1(e14, e13)) = op2(h14(e14), h14(e13))) & (h14(op1(e14, e12)) = op2(h14(e14), h14(e12))) & (h14(op1(e14, e11)) = op2(h14(e14), h14(e11))) & (h14(op1(e14, e10)) = op2(h14(e14), h14(e10))) & (h14(op1(e13, e14)) = op2(h14(e13), h14(e14))) & (h14(op1(e13, e13)) = op2(h14(e13), h14(e13))) & (h14(op1(e13, e12)) = op2(h14(e13), h14(e12))) & (h14(op1(e13, e11)) = op2(h14(e13), h14(e11))) & (h14(op1(e13, e10)) = op2(h14(e13), h14(e10))) & (h14(op1(e12, e14)) = op2(h14(e12), h14(e14))) & (h14(op1(e12, e13)) = op2(h14(e12), h14(e13))) & (h14(op1(e12, e12)) = op2(h14(e12), h14(e12))) & (h14(op1(e12, e11)) = op2(h14(e12), h14(e11))) & (h14(op1(e12, e10)) = op2(h14(e12), h14(e10))) & (h14(op1(e11, e14)) = op2(h14(e11), h14(e14))) & (h14(op1(e11, e13)) = op2(h14(e11), h14(e13))) & (h14(op1(e11, e12)) = op2(h14(e11), h14(e12))) & (h14(op1(e11, e11)) = op2(h14(e11), h14(e11))) & (h14(op1(e11, e10)) = op2(h14(e11), h14(e10))) & (h14(op1(e10, e14)) = op2(h14(e10), h14(e14))) & (h14(op1(e10, e13)) = op2(h14(e10), h14(e13))) & (h14(op1(e10, e12)) = op2(h14(e10), h14(e12))) & (h14(op1(e10, e11)) = op2(h14(e10), h14(e11))) & (h14(op1(e10, e10)) = op2(h14(e10), h14(e10)))) | (((e24 = h13(e14)) | (e24 = h13(e13)) | (e24 = h13(e12)) | (e24 = h13(e11)) | (e24 = h13(e10))) & ((e23 = h13(e14)) | (e23 = h13(e13)) | (e23 = h13(e12)) | (e23 = h13(e11)) | (e23 = h13(e10))) & ((e22 = h13(e14)) | (e22 = h13(e13)) | (e22 = h13(e12)) | (e22 = h13(e11)) | (e22 = h13(e10))) & ((e21 = h13(e14)) | (e21 = h13(e13)) | (e21 = h13(e12)) | (e21 = h13(e11)) | (e21 = h13(e10))) & ((e20 = h13(e14)) | (e20 = h13(e13)) | (e20 = h13(e12)) | (e20 = h13(e11)) | (e20 = h13(e10))) & (h13(op1(e14, e14)) = op2(h13(e14), h13(e14))) & (h13(op1(e14, e13)) = op2(h13(e14), h13(e13))) & (h13(op1(e14, e12)) = op2(h13(e14), h13(e12))) & (h13(op1(e14, e11)) = op2(h13(e14), h13(e11))) & (h13(op1(e14, e10)) = op2(h13(e14), h13(e10))) & (h13(op1(e13, e14)) = op2(h13(e13), h13(e14))) & (h13(op1(e13, e13)) = op2(h13(e13), h13(e13))) & (h13(op1(e13, e12)) = op2(h13(e13), h13(e12))) & (h13(op1(e13, e11)) = op2(h13(e13), h13(e11))) & (h13(op1(e13, e10)) = op2(h13(e13), h13(e10))) & (h13(op1(e12, e14)) = op2(h13(e12), h13(e14))) & (h13(op1(e12, e13)) = op2(h13(e12), h13(e13))) & (h13(op1(e12, e12)) = op2(h13(e12), h13(e12))) & (h13(op1(e12, e11)) = op2(h13(e12), h13(e11))) & (h13(op1(e12, e10)) = op2(h13(e12), h13(e10))) & (h13(op1(e11, e14)) = op2(h13(e11), h13(e14))) & (h13(op1(e11, e13)) = op2(h13(e11), h13(e13))) & (h13(op1(e11, e12)) = op2(h13(e11), h13(e12))) & (h13(op1(e11, e11)) = op2(h13(e11), h13(e11))) & (h13(op1(e11, e10)) = op2(h13(e11), h13(e10))) & (h13(op1(e10, e14)) = op2(h13(e10), h13(e14))) & (h13(op1(e10, e13)) = op2(h13(e10), h13(e13))) & (h13(op1(e10, e12)) = op2(h13(e10), h13(e12))) & (h13(op1(e10, e11)) = op2(h13(e10), h13(e11))) & (h13(op1(e10, e10)) = op2(h13(e10), h13(e10)))) | (((e24 = h12(e14)) | (e24 = h12(e13)) | (e24 = h12(e12)) | (e24 = h12(e11)) | (e24 = h12(e10))) & ((e23 = h12(e14)) | (e23 = h12(e13)) | (e23 = h12(e12)) | (e23 = h12(e11)) | (e23 = h12(e10))) & ((e22 = h12(e14)) | (e22 = h12(e13)) | (e22 = h12(e12)) | (e22 = h12(e11)) | (e22 = h12(e10))) & ((e21 = h12(e14)) | (e21 = h12(e13)) | (e21 = h12(e12)) | (e21 = h12(e11)) | (e21 = h12(e10))) & ((e20 = h12(e14)) | (e20 = h12(e13)) | (e20 = h12(e12)) | (e20 = h12(e11)) | (e20 = h12(e10))) & (h12(op1(e14, e14)) = op2(h12(e14), h12(e14))) & (h12(op1(e14, e13)) = op2(h12(e14), h12(e13))) & (h12(op1(e14, e12)) = op2(h12(e14), h12(e12))) & (h12(op1(e14, e11)) = op2(h12(e14), h12(e11))) & (h12(op1(e14, e10)) = op2(h12(e14), h12(e10))) & (h12(op1(e13, e14)) = op2(h12(e13), h12(e14))) & (h12(op1(e13, e13)) = op2(h12(e13), h12(e13))) & (h12(op1(e13, e12)) = op2(h12(e13), h12(e12))) & (h12(op1(e13, e11)) = op2(h12(e13), h12(e11))) & (h12(op1(e13, e10)) = op2(h12(e13), h12(e10))) & (h12(op1(e12, e14)) = op2(h12(e12), h12(e14))) & (h12(op1(e12, e13)) = op2(h12(e12), h12(e13))) & (h12(op1(e12, e12)) = op2(h12(e12), h12(e12))) & (h12(op1(e12, e11)) = op2(h12(e12), h12(e11))) & (h12(op1(e12, e10)) = op2(h12(e12), h12(e10))) & (h12(op1(e11, e14)) = op2(h12(e11), h12(e14))) & (h12(op1(e11, e13)) = op2(h12(e11), h12(e13))) & (h12(op1(e11, e12)) = op2(h12(e11), h12(e12))) & (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) & (h12(op1(e11, e10)) = op2(h12(e11), h12(e10))) & (h12(op1(e10, e14)) = op2(h12(e10), h12(e14))) & (h12(op1(e10, e13)) = op2(h12(e10), h12(e13))) & (h12(op1(e10, e12)) = op2(h12(e10), h12(e12))) & (h12(op1(e10, e11)) = op2(h12(e10), h12(e11))) & (h12(op1(e10, e10)) = op2(h12(e10), h12(e10)))) | (((e24 = h11(e14)) | (e24 = h11(e13)) | (e24 = h11(e12)) | (e24 = h11(e11)) | (e24 = h11(e10))) & ((e23 = h11(e14)) | (e23 = h11(e13)) | (e23 = h11(e12)) | (e23 = h11(e11)) | (e23 = h11(e10))) & ((e22 = h11(e14)) | (e22 = h11(e13)) | (e22 = h11(e12)) | (e22 = h11(e11)) | (e22 = h11(e10))) & ((e21 = h11(e14)) | (e21 = h11(e13)) | (e21 = h11(e12)) | (e21 = h11(e11)) | (e21 = h11(e10))) & ((e20 = h11(e14)) | (e20 = h11(e13)) | (e20 = h11(e12)) | (e20 = h11(e11)) | (e20 = h11(e10))) & (h11(op1(e14, e14)) = op2(h11(e14), h11(e14))) & (h11(op1(e14, e13)) = op2(h11(e14), h11(e13))) & (h11(op1(e14, e12)) = op2(h11(e14), h11(e12))) & (h11(op1(e14, e11)) = op2(h11(e14), h11(e11))) & (h11(op1(e14, e10)) = op2(h11(e14), h11(e10))) & (h11(op1(e13, e14)) = op2(h11(e13), h11(e14))) & (h11(op1(e13, e13)) = op2(h11(e13), h11(e13))) & (h11(op1(e13, e12)) = op2(h11(e13), h11(e12))) & (h11(op1(e13, e11)) = op2(h11(e13), h11(e11))) & (h11(op1(e13, e10)) = op2(h11(e13), h11(e10))) & (h11(op1(e12, e14)) = op2(h11(e12), h11(e14))) & (h11(op1(e12, e13)) = op2(h11(e12), h11(e13))) & (h11(op1(e12, e12)) = op2(h11(e12), h11(e12))) & (h11(op1(e12, e11)) = op2(h11(e12), h11(e11))) & (h11(op1(e12, e10)) = op2(h11(e12), h11(e10))) & (h11(op1(e11, e14)) = op2(h11(e11), h11(e14))) & (h11(op1(e11, e13)) = op2(h11(e11), h11(e13))) & (h11(op1(e11, e12)) = op2(h11(e11), h11(e12))) & (h11(op1(e11, e11)) = op2(h11(e11), h11(e11))) & (h11(op1(e11, e10)) = op2(h11(e11), h11(e10))) & (h11(op1(e10, e14)) = op2(h11(e10), h11(e14))) & (h11(op1(e10, e13)) = op2(h11(e10), h11(e13))) & (h11(op1(e10, e12)) = op2(h11(e10), h11(e12))) & (h11(op1(e10, e11)) = op2(h11(e10), h11(e11))) & (h11(op1(e10, e10)) = op2(h11(e10), h11(e10)))) | (((e24 = h10(e14)) | (e24 = h10(e13)) | (e24 = h10(e12)) | (e24 = h10(e11)) | (e24 = h10(e10))) & ((e23 = h10(e14)) | (e23 = h10(e13)) | (e23 = h10(e12)) | (e23 = h10(e11)) | (e23 = h10(e10))) & ((e22 = h10(e14)) | (e22 = h10(e13)) | (e22 = h10(e12)) | (e22 = h10(e11)) | (e22 = h10(e10))) & ((e21 = h10(e14)) | (e21 = h10(e13)) | (e21 = h10(e12)) | (e21 = h10(e11)) | (e21 = h10(e10))) & ((e20 = h10(e14)) | (e20 = h10(e13)) | (e20 = h10(e12)) | (e20 = h10(e11)) | (e20 = h10(e10))) & (h10(op1(e14, e14)) = op2(h10(e14), h10(e14))) & (h10(op1(e14, e13)) = op2(h10(e14), h10(e13))) & (h10(op1(e14, e12)) = op2(h10(e14), h10(e12))) & (h10(op1(e14, e11)) = op2(h10(e14), h10(e11))) & (h10(op1(e14, e10)) = op2(h10(e14), h10(e10))) & (h10(op1(e13, e14)) = op2(h10(e13), h10(e14))) & (h10(op1(e13, e13)) = op2(h10(e13), h10(e13))) & (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) & (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) & (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) & (h10(op1(e12, e14)) = op2(h10(e12), h10(e14))) & (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) & (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) & (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) & (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) & (h10(op1(e11, e14)) = op2(h10(e11), h10(e14))) & (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) & (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) & (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) & (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) & (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) & (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) & (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) & (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) & (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))) | (((e24 = h9(e14)) | (e24 = h9(e13)) | (e24 = h9(e12)) | (e24 = h9(e11)) | (e24 = h9(e10))) & ((e23 = h9(e14)) | (e23 = h9(e13)) | (e23 = h9(e12)) | (e23 = h9(e11)) | (e23 = h9(e10))) & ((e22 = h9(e14)) | (e22 = h9(e13)) | (e22 = h9(e12)) | (e22 = h9(e11)) | (e22 = h9(e10))) & ((e21 = h9(e14)) | (e21 = h9(e13)) | (e21 = h9(e12)) | (e21 = h9(e11)) | (e21 = h9(e10))) & ((e20 = h9(e14)) | (e20 = h9(e13)) | (e20 = h9(e12)) | (e20 = h9(e11)) | (e20 = h9(e10))) & (h9(op1(e14, e14)) = op2(h9(e14), h9(e14))) & (h9(op1(e14, e13)) = op2(h9(e14), h9(e13))) & (h9(op1(e14, e12)) = op2(h9(e14), h9(e12))) & (h9(op1(e14, e11)) = op2(h9(e14), h9(e11))) & (h9(op1(e14, e10)) = op2(h9(e14), h9(e10))) & (h9(op1(e13, e14)) = op2(h9(e13), h9(e14))) & (h9(op1(e13, e13)) = op2(h9(e13), h9(e13))) & (h9(op1(e13, e12)) = op2(h9(e13), h9(e12))) & (h9(op1(e13, e11)) = op2(h9(e13), h9(e11))) & (h9(op1(e13, e10)) = op2(h9(e13), h9(e10))) & (h9(op1(e12, e14)) = op2(h9(e12), h9(e14))) & (h9(op1(e12, e13)) = op2(h9(e12), h9(e13))) & (h9(op1(e12, e12)) = op2(h9(e12), h9(e12))) & (h9(op1(e12, e11)) = op2(h9(e12), h9(e11))) & (h9(op1(e12, e10)) = op2(h9(e12), h9(e10))) & (h9(op1(e11, e14)) = op2(h9(e11), h9(e14))) & (h9(op1(e11, e13)) = op2(h9(e11), h9(e13))) & (h9(op1(e11, e12)) = op2(h9(e11), h9(e12))) & (h9(op1(e11, e11)) = op2(h9(e11), h9(e11))) & (h9(op1(e11, e10)) = op2(h9(e11), h9(e10))) & (h9(op1(e10, e14)) = op2(h9(e10), h9(e14))) & (h9(op1(e10, e13)) = op2(h9(e10), h9(e13))) & (h9(op1(e10, e12)) = op2(h9(e10), h9(e12))) & (h9(op1(e10, e11)) = op2(h9(e10), h9(e11))) & (h9(op1(e10, e10)) = op2(h9(e10), h9(e10)))) | (((e24 = h8(e14)) | (e24 = h8(e13)) | (e24 = h8(e12)) | (e24 = h8(e11)) | (e24 = h8(e10))) & ((e23 = h8(e14)) | (e23 = h8(e13)) | (e23 = h8(e12)) | (e23 = h8(e11)) | (e23 = h8(e10))) & ((e22 = h8(e14)) | (e22 = h8(e13)) | (e22 = h8(e12)) | (e22 = h8(e11)) | (e22 = h8(e10))) & ((e21 = h8(e14)) | (e21 = h8(e13)) | (e21 = h8(e12)) | (e21 = h8(e11)) | (e21 = h8(e10))) & ((e20 = h8(e14)) | (e20 = h8(e13)) | (e20 = h8(e12)) | (e20 = h8(e11)) | (e20 = h8(e10))) & (h8(op1(e14, e14)) = op2(h8(e14), h8(e14))) & (h8(op1(e14, e13)) = op2(h8(e14), h8(e13))) & (h8(op1(e14, e12)) = op2(h8(e14), h8(e12))) & (h8(op1(e14, e11)) = op2(h8(e14), h8(e11))) & (h8(op1(e14, e10)) = op2(h8(e14), h8(e10))) & (h8(op1(e13, e14)) = op2(h8(e13), h8(e14))) & (h8(op1(e13, e13)) = op2(h8(e13), h8(e13))) & (h8(op1(e13, e12)) = op2(h8(e13), h8(e12))) & (h8(op1(e13, e11)) = op2(h8(e13), h8(e11))) & (h8(op1(e13, e10)) = op2(h8(e13), h8(e10))) & (h8(op1(e12, e14)) = op2(h8(e12), h8(e14))) & (h8(op1(e12, e13)) = op2(h8(e12), h8(e13))) & (h8(op1(e12, e12)) = op2(h8(e12), h8(e12))) & (h8(op1(e12, e11)) = op2(h8(e12), h8(e11))) & (h8(op1(e12, e10)) = op2(h8(e12), h8(e10))) & (h8(op1(e11, e14)) = op2(h8(e11), h8(e14))) & (h8(op1(e11, e13)) = op2(h8(e11), h8(e13))) & (h8(op1(e11, e12)) = op2(h8(e11), h8(e12))) & (h8(op1(e11, e11)) = op2(h8(e11), h8(e11))) & (h8(op1(e11, e10)) = op2(h8(e11), h8(e10))) & (h8(op1(e10, e14)) = op2(h8(e10), h8(e14))) & (h8(op1(e10, e13)) = op2(h8(e10), h8(e13))) & (h8(op1(e10, e12)) = op2(h8(e10), h8(e12))) & (h8(op1(e10, e11)) = op2(h8(e10), h8(e11))) & (h8(op1(e10, e10)) = op2(h8(e10), h8(e10)))) | (((e24 = h7(e14)) | (e24 = h7(e13)) | (e24 = h7(e12)) | (e24 = h7(e11)) | (e24 = h7(e10))) & ((e23 = h7(e14)) | (e23 = h7(e13)) | (e23 = h7(e12)) | (e23 = h7(e11)) | (e23 = h7(e10))) & ((e22 = h7(e14)) | (e22 = h7(e13)) | (e22 = h7(e12)) | (e22 = h7(e11)) | (e22 = h7(e10))) & ((e21 = h7(e14)) | (e21 = h7(e13)) | (e21 = h7(e12)) | (e21 = h7(e11)) | (e21 = h7(e10))) & ((e20 = h7(e14)) | (e20 = h7(e13)) | (e20 = h7(e12)) | (e20 = h7(e11)) | (e20 = h7(e10))) & (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) & (h7(op1(e14, e13)) = op2(h7(e14), h7(e13))) & (h7(op1(e14, e12)) = op2(h7(e14), h7(e12))) & (h7(op1(e14, e11)) = op2(h7(e14), h7(e11))) & (h7(op1(e14, e10)) = op2(h7(e14), h7(e10))) & (h7(op1(e13, e14)) = op2(h7(e13), h7(e14))) & (h7(op1(e13, e13)) = op2(h7(e13), h7(e13))) & (h7(op1(e13, e12)) = op2(h7(e13), h7(e12))) & (h7(op1(e13, e11)) = op2(h7(e13), h7(e11))) & (h7(op1(e13, e10)) = op2(h7(e13), h7(e10))) & (h7(op1(e12, e14)) = op2(h7(e12), h7(e14))) & (h7(op1(e12, e13)) = op2(h7(e12), h7(e13))) & (h7(op1(e12, e12)) = op2(h7(e12), h7(e12))) & (h7(op1(e12, e11)) = op2(h7(e12), h7(e11))) & (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) & (h7(op1(e11, e14)) = op2(h7(e11), h7(e14))) & (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) & (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) & (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) & (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) & (h7(op1(e10, e14)) = op2(h7(e10), h7(e14))) & (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) & (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) & (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) & (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))) | (((e24 = h6(e14)) | (e24 = h6(e13)) | (e24 = h6(e12)) | (e24 = h6(e11)) | (e24 = h6(e10))) & ((e23 = h6(e14)) | (e23 = h6(e13)) | (e23 = h6(e12)) | (e23 = h6(e11)) | (e23 = h6(e10))) & ((e22 = h6(e14)) | (e22 = h6(e13)) | (e22 = h6(e12)) | (e22 = h6(e11)) | (e22 = h6(e10))) & ((e21 = h6(e14)) | (e21 = h6(e13)) | (e21 = h6(e12)) | (e21 = h6(e11)) | (e21 = h6(e10))) & ((e20 = h6(e14)) | (e20 = h6(e13)) | (e20 = h6(e12)) | (e20 = h6(e11)) | (e20 = h6(e10))) & (h6(op1(e14, e14)) = op2(h6(e14), h6(e14))) & (h6(op1(e14, e13)) = op2(h6(e14), h6(e13))) & (h6(op1(e14, e12)) = op2(h6(e14), h6(e12))) & (h6(op1(e14, e11)) = op2(h6(e14), h6(e11))) & (h6(op1(e14, e10)) = op2(h6(e14), h6(e10))) & (h6(op1(e13, e14)) = op2(h6(e13), h6(e14))) & (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) & (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) & (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) & (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) & (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) & (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) & (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) & (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) & (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) & (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) & (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) & (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) & (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) & (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) & (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) & (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) & (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) & (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) & (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))) | (((e24 = h5(e14)) | (e24 = h5(e13)) | (e24 = h5(e12)) | (e24 = h5(e11)) | (e24 = h5(e10))) & ((e23 = h5(e14)) | (e23 = h5(e13)) | (e23 = h5(e12)) | (e23 = h5(e11)) | (e23 = h5(e10))) & ((e22 = h5(e14)) | (e22 = h5(e13)) | (e22 = h5(e12)) | (e22 = h5(e11)) | (e22 = h5(e10))) & ((e21 = h5(e14)) | (e21 = h5(e13)) | (e21 = h5(e12)) | (e21 = h5(e11)) | (e21 = h5(e10))) & ((e20 = h5(e14)) | (e20 = h5(e13)) | (e20 = h5(e12)) | (e20 = h5(e11)) | (e20 = h5(e10))) & (h5(op1(e14, e14)) = op2(h5(e14), h5(e14))) & (h5(op1(e14, e13)) = op2(h5(e14), h5(e13))) & (h5(op1(e14, e12)) = op2(h5(e14), h5(e12))) & (h5(op1(e14, e11)) = op2(h5(e14), h5(e11))) & (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) & (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) & (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) & (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) & (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) & (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) & (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) & (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) & (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) & (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) & (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) & (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) & (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) & (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) & (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) & (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) & (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) & (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) & (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) & (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) & (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))) | (((e24 = h4(e14)) | (e24 = h4(e13)) | (e24 = h4(e12)) | (e24 = h4(e11)) | (e24 = h4(e10))) & ((e23 = h4(e14)) | (e23 = h4(e13)) | (e23 = h4(e12)) | (e23 = h4(e11)) | (e23 = h4(e10))) & ((e22 = h4(e14)) | (e22 = h4(e13)) | (e22 = h4(e12)) | (e22 = h4(e11)) | (e22 = h4(e10))) & ((e21 = h4(e14)) | (e21 = h4(e13)) | (e21 = h4(e12)) | (e21 = h4(e11)) | (e21 = h4(e10))) & ((e20 = h4(e14)) | (e20 = h4(e13)) | (e20 = h4(e12)) | (e20 = h4(e11)) | (e20 = h4(e10))) & (h4(op1(e14, e14)) = op2(h4(e14), h4(e14))) & (h4(op1(e14, e13)) = op2(h4(e14), h4(e13))) & (h4(op1(e14, e12)) = op2(h4(e14), h4(e12))) & (h4(op1(e14, e11)) = op2(h4(e14), h4(e11))) & (h4(op1(e14, e10)) = op2(h4(e14), h4(e10))) & (h4(op1(e13, e14)) = op2(h4(e13), h4(e14))) & (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) & (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) & (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) & (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) & (h4(op1(e12, e14)) = op2(h4(e12), h4(e14))) & (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) & (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) & (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) & (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) & (h4(op1(e11, e14)) = op2(h4(e11), h4(e14))) & (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) & (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) & (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) & (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) & (h4(op1(e10, e14)) = op2(h4(e10), h4(e14))) & (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) & (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) & (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) & (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) | (((e24 = h3(e14)) | (e24 = h3(e13)) | (e24 = h3(e12)) | (e24 = h3(e11)) | (e24 = h3(e10))) & ((e23 = h3(e14)) | (e23 = h3(e13)) | (e23 = h3(e12)) | (e23 = h3(e11)) | (e23 = h3(e10))) & ((e22 = h3(e14)) | (e22 = h3(e13)) | (e22 = h3(e12)) | (e22 = h3(e11)) | (e22 = h3(e10))) & ((e21 = h3(e14)) | (e21 = h3(e13)) | (e21 = h3(e12)) | (e21 = h3(e11)) | (e21 = h3(e10))) & ((e20 = h3(e14)) | (e20 = h3(e13)) | (e20 = h3(e12)) | (e20 = h3(e11)) | (e20 = h3(e10))) & (h3(op1(e14, e14)) = op2(h3(e14), h3(e14))) & (h3(op1(e14, e13)) = op2(h3(e14), h3(e13))) & (h3(op1(e14, e12)) = op2(h3(e14), h3(e12))) & (h3(op1(e14, e11)) = op2(h3(e14), h3(e11))) & (h3(op1(e14, e10)) = op2(h3(e14), h3(e10))) & (h3(op1(e13, e14)) = op2(h3(e13), h3(e14))) & (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) & (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) & (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) & (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) & (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) & (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) & (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) & (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) & (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) & (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) & (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) & (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) & (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) & (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) & (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) & (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) & (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) & (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) & (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) | (((e24 = h2(e14)) | (e24 = h2(e13)) | (e24 = h2(e12)) | (e24 = h2(e11)) | (e24 = h2(e10))) & ((e23 = h2(e14)) | (e23 = h2(e13)) | (e23 = h2(e12)) | (e23 = h2(e11)) | (e23 = h2(e10))) & ((e22 = h2(e14)) | (e22 = h2(e13)) | (e22 = h2(e12)) | (e22 = h2(e11)) | (e22 = h2(e10))) & ((e21 = h2(e14)) | (e21 = h2(e13)) | (e21 = h2(e12)) | (e21 = h2(e11)) | (e21 = h2(e10))) & ((e20 = h2(e14)) | (e20 = h2(e13)) | (e20 = h2(e12)) | (e20 = h2(e11)) | (e20 = h2(e10))) & (h2(op1(e14, e14)) = op2(h2(e14), h2(e14))) & (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) & (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) & (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) & (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) & (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) & (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) & (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) & (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) & (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) & (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) & (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) & (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) & (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) & (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) & (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) & (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) & (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) & (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) & (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) & (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) & (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) & (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) & (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) & (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) | (((e24 = h1(e14)) | (e24 = h1(e13)) | (e24 = h1(e12)) | (e24 = h1(e11)) | (e24 = h1(e10))) & ((e23 = h1(e14)) | (e23 = h1(e13)) | (e23 = h1(e12)) | (e23 = h1(e11)) | (e23 = h1(e10))) & ((e22 = h1(e14)) | (e22 = h1(e13)) | (e22 = h1(e12)) | (e22 = h1(e11)) | (e22 = h1(e10))) & ((e21 = h1(e14)) | (e21 = h1(e13)) | (e21 = h1(e12)) | (e21 = h1(e11)) | (e21 = h1(e10))) & ((e20 = h1(e14)) | (e20 = h1(e13)) | (e20 = h1(e12)) | (e20 = h1(e11)) | (e20 = h1(e10))) & (h1(op1(e14, e14)) = op2(h1(e14), h1(e14))) & (h1(op1(e14, e13)) = op2(h1(e14), h1(e13))) & (h1(op1(e14, e12)) = op2(h1(e14), h1(e12))) & (h1(op1(e14, e11)) = op2(h1(e14), h1(e11))) & (h1(op1(e14, e10)) = op2(h1(e14), h1(e10))) & (h1(op1(e13, e14)) = op2(h1(e13), h1(e14))) & (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) & (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) & (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) & (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) & (h1(op1(e12, e14)) = op2(h1(e12), h1(e14))) & (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) & (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) & (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) & (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) & (h1(op1(e11, e14)) = op2(h1(e11), h1(e14))) & (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) & (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) & (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) & (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) & (h1(op1(e10, e14)) = op2(h1(e10), h1(e14))) & (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) & (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) & (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) & (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(negated_conjecture, [], [f36])).
fof(f36, plain, ~ ((((e24 = h20(e14)) | (e24 = h20(e13)) | (e24 = h20(e12)) | (e24 = h20(e11)) | (e24 = h20(e10))) & ((e23 = h20(e14)) | (e23 = h20(e13)) | (e23 = h20(e12)) | (e23 = h20(e11)) | (e23 = h20(e10))) & ((e22 = h20(e14)) | (e22 = h20(e13)) | (e22 = h20(e12)) | (e22 = h20(e11)) | (e22 = h20(e10))) & ((e21 = h20(e14)) | (e21 = h20(e13)) | (e21 = h20(e12)) | (e21 = h20(e11)) | (e21 = h20(e10))) & ((e20 = h20(e14)) | (e20 = h20(e13)) | (e20 = h20(e12)) | (e20 = h20(e11)) | (e20 = h20(e10))) & (h20(op1(e14, e14)) = op2(h20(e14), h20(e14))) & (h20(op1(e14, e13)) = op2(h20(e14), h20(e13))) & (h20(op1(e14, e12)) = op2(h20(e14), h20(e12))) & (h20(op1(e14, e11)) = op2(h20(e14), h20(e11))) & (h20(op1(e14, e10)) = op2(h20(e14), h20(e10))) & (h20(op1(e13, e14)) = op2(h20(e13), h20(e14))) & (h20(op1(e13, e13)) = op2(h20(e13), h20(e13))) & (h20(op1(e13, e12)) = op2(h20(e13), h20(e12))) & (h20(op1(e13, e11)) = op2(h20(e13), h20(e11))) & (h20(op1(e13, e10)) = op2(h20(e13), h20(e10))) & (h20(op1(e12, e14)) = op2(h20(e12), h20(e14))) & (h20(op1(e12, e13)) = op2(h20(e12), h20(e13))) & (h20(op1(e12, e12)) = op2(h20(e12), h20(e12))) & (h20(op1(e12, e11)) = op2(h20(e12), h20(e11))) & (h20(op1(e12, e10)) = op2(h20(e12), h20(e10))) & (h20(op1(e11, e14)) = op2(h20(e11), h20(e14))) & (h20(op1(e11, e13)) = op2(h20(e11), h20(e13))) & (h20(op1(e11, e12)) = op2(h20(e11), h20(e12))) & (h20(op1(e11, e11)) = op2(h20(e11), h20(e11))) & (h20(op1(e11, e10)) = op2(h20(e11), h20(e10))) & (h20(op1(e10, e14)) = op2(h20(e10), h20(e14))) & (h20(op1(e10, e13)) = op2(h20(e10), h20(e13))) & (h20(op1(e10, e12)) = op2(h20(e10), h20(e12))) & (h20(op1(e10, e11)) = op2(h20(e10), h20(e11))) & (h20(op1(e10, e10)) = op2(h20(e10), h20(e10)))) | (((e24 = h19(e14)) | (e24 = h19(e13)) | (e24 = h19(e12)) | (e24 = h19(e11)) | (e24 = h19(e10))) & ((e23 = h19(e14)) | (e23 = h19(e13)) | (e23 = h19(e12)) | (e23 = h19(e11)) | (e23 = h19(e10))) & ((e22 = h19(e14)) | (e22 = h19(e13)) | (e22 = h19(e12)) | (e22 = h19(e11)) | (e22 = h19(e10))) & ((e21 = h19(e14)) | (e21 = h19(e13)) | (e21 = h19(e12)) | (e21 = h19(e11)) | (e21 = h19(e10))) & ((e20 = h19(e14)) | (e20 = h19(e13)) | (e20 = h19(e12)) | (e20 = h19(e11)) | (e20 = h19(e10))) & (h19(op1(e14, e14)) = op2(h19(e14), h19(e14))) & (h19(op1(e14, e13)) = op2(h19(e14), h19(e13))) & (h19(op1(e14, e12)) = op2(h19(e14), h19(e12))) & (h19(op1(e14, e11)) = op2(h19(e14), h19(e11))) & (h19(op1(e14, e10)) = op2(h19(e14), h19(e10))) & (h19(op1(e13, e14)) = op2(h19(e13), h19(e14))) & (h19(op1(e13, e13)) = op2(h19(e13), h19(e13))) & (h19(op1(e13, e12)) = op2(h19(e13), h19(e12))) & (h19(op1(e13, e11)) = op2(h19(e13), h19(e11))) & (h19(op1(e13, e10)) = op2(h19(e13), h19(e10))) & (h19(op1(e12, e14)) = op2(h19(e12), h19(e14))) & (h19(op1(e12, e13)) = op2(h19(e12), h19(e13))) & (h19(op1(e12, e12)) = op2(h19(e12), h19(e12))) & (h19(op1(e12, e11)) = op2(h19(e12), h19(e11))) & (h19(op1(e12, e10)) = op2(h19(e12), h19(e10))) & (h19(op1(e11, e14)) = op2(h19(e11), h19(e14))) & (h19(op1(e11, e13)) = op2(h19(e11), h19(e13))) & (h19(op1(e11, e12)) = op2(h19(e11), h19(e12))) & (h19(op1(e11, e11)) = op2(h19(e11), h19(e11))) & (h19(op1(e11, e10)) = op2(h19(e11), h19(e10))) & (h19(op1(e10, e14)) = op2(h19(e10), h19(e14))) & (h19(op1(e10, e13)) = op2(h19(e10), h19(e13))) & (h19(op1(e10, e12)) = op2(h19(e10), h19(e12))) & (h19(op1(e10, e11)) = op2(h19(e10), h19(e11))) & (h19(op1(e10, e10)) = op2(h19(e10), h19(e10)))) | (((e24 = h18(e14)) | (e24 = h18(e13)) | (e24 = h18(e12)) | (e24 = h18(e11)) | (e24 = h18(e10))) & ((e23 = h18(e14)) | (e23 = h18(e13)) | (e23 = h18(e12)) | (e23 = h18(e11)) | (e23 = h18(e10))) & ((e22 = h18(e14)) | (e22 = h18(e13)) | (e22 = h18(e12)) | (e22 = h18(e11)) | (e22 = h18(e10))) & ((e21 = h18(e14)) | (e21 = h18(e13)) | (e21 = h18(e12)) | (e21 = h18(e11)) | (e21 = h18(e10))) & ((e20 = h18(e14)) | (e20 = h18(e13)) | (e20 = h18(e12)) | (e20 = h18(e11)) | (e20 = h18(e10))) & (h18(op1(e14, e14)) = op2(h18(e14), h18(e14))) & (h18(op1(e14, e13)) = op2(h18(e14), h18(e13))) & (h18(op1(e14, e12)) = op2(h18(e14), h18(e12))) & (h18(op1(e14, e11)) = op2(h18(e14), h18(e11))) & (h18(op1(e14, e10)) = op2(h18(e14), h18(e10))) & (h18(op1(e13, e14)) = op2(h18(e13), h18(e14))) & (h18(op1(e13, e13)) = op2(h18(e13), h18(e13))) & (h18(op1(e13, e12)) = op2(h18(e13), h18(e12))) & (h18(op1(e13, e11)) = op2(h18(e13), h18(e11))) & (h18(op1(e13, e10)) = op2(h18(e13), h18(e10))) & (h18(op1(e12, e14)) = op2(h18(e12), h18(e14))) & (h18(op1(e12, e13)) = op2(h18(e12), h18(e13))) & (h18(op1(e12, e12)) = op2(h18(e12), h18(e12))) & (h18(op1(e12, e11)) = op2(h18(e12), h18(e11))) & (h18(op1(e12, e10)) = op2(h18(e12), h18(e10))) & (h18(op1(e11, e14)) = op2(h18(e11), h18(e14))) & (h18(op1(e11, e13)) = op2(h18(e11), h18(e13))) & (h18(op1(e11, e12)) = op2(h18(e11), h18(e12))) & (h18(op1(e11, e11)) = op2(h18(e11), h18(e11))) & (h18(op1(e11, e10)) = op2(h18(e11), h18(e10))) & (h18(op1(e10, e14)) = op2(h18(e10), h18(e14))) & (h18(op1(e10, e13)) = op2(h18(e10), h18(e13))) & (h18(op1(e10, e12)) = op2(h18(e10), h18(e12))) & (h18(op1(e10, e11)) = op2(h18(e10), h18(e11))) & (h18(op1(e10, e10)) = op2(h18(e10), h18(e10)))) | (((e24 = h17(e14)) | (e24 = h17(e13)) | (e24 = h17(e12)) | (e24 = h17(e11)) | (e24 = h17(e10))) & ((e23 = h17(e14)) | (e23 = h17(e13)) | (e23 = h17(e12)) | (e23 = h17(e11)) | (e23 = h17(e10))) & ((e22 = h17(e14)) | (e22 = h17(e13)) | (e22 = h17(e12)) | (e22 = h17(e11)) | (e22 = h17(e10))) & ((e21 = h17(e14)) | (e21 = h17(e13)) | (e21 = h17(e12)) | (e21 = h17(e11)) | (e21 = h17(e10))) & ((e20 = h17(e14)) | (e20 = h17(e13)) | (e20 = h17(e12)) | (e20 = h17(e11)) | (e20 = h17(e10))) & (h17(op1(e14, e14)) = op2(h17(e14), h17(e14))) & (h17(op1(e14, e13)) = op2(h17(e14), h17(e13))) & (h17(op1(e14, e12)) = op2(h17(e14), h17(e12))) & (h17(op1(e14, e11)) = op2(h17(e14), h17(e11))) & (h17(op1(e14, e10)) = op2(h17(e14), h17(e10))) & (h17(op1(e13, e14)) = op2(h17(e13), h17(e14))) & (h17(op1(e13, e13)) = op2(h17(e13), h17(e13))) & (h17(op1(e13, e12)) = op2(h17(e13), h17(e12))) & (h17(op1(e13, e11)) = op2(h17(e13), h17(e11))) & (h17(op1(e13, e10)) = op2(h17(e13), h17(e10))) & (h17(op1(e12, e14)) = op2(h17(e12), h17(e14))) & (h17(op1(e12, e13)) = op2(h17(e12), h17(e13))) & (h17(op1(e12, e12)) = op2(h17(e12), h17(e12))) & (h17(op1(e12, e11)) = op2(h17(e12), h17(e11))) & (h17(op1(e12, e10)) = op2(h17(e12), h17(e10))) & (h17(op1(e11, e14)) = op2(h17(e11), h17(e14))) & (h17(op1(e11, e13)) = op2(h17(e11), h17(e13))) & (h17(op1(e11, e12)) = op2(h17(e11), h17(e12))) & (h17(op1(e11, e11)) = op2(h17(e11), h17(e11))) & (h17(op1(e11, e10)) = op2(h17(e11), h17(e10))) & (h17(op1(e10, e14)) = op2(h17(e10), h17(e14))) & (h17(op1(e10, e13)) = op2(h17(e10), h17(e13))) & (h17(op1(e10, e12)) = op2(h17(e10), h17(e12))) & (h17(op1(e10, e11)) = op2(h17(e10), h17(e11))) & (h17(op1(e10, e10)) = op2(h17(e10), h17(e10)))) | (((e24 = h16(e14)) | (e24 = h16(e13)) | (e24 = h16(e12)) | (e24 = h16(e11)) | (e24 = h16(e10))) & ((e23 = h16(e14)) | (e23 = h16(e13)) | (e23 = h16(e12)) | (e23 = h16(e11)) | (e23 = h16(e10))) & ((e22 = h16(e14)) | (e22 = h16(e13)) | (e22 = h16(e12)) | (e22 = h16(e11)) | (e22 = h16(e10))) & ((e21 = h16(e14)) | (e21 = h16(e13)) | (e21 = h16(e12)) | (e21 = h16(e11)) | (e21 = h16(e10))) & ((e20 = h16(e14)) | (e20 = h16(e13)) | (e20 = h16(e12)) | (e20 = h16(e11)) | (e20 = h16(e10))) & (h16(op1(e14, e14)) = op2(h16(e14), h16(e14))) & (h16(op1(e14, e13)) = op2(h16(e14), h16(e13))) & (h16(op1(e14, e12)) = op2(h16(e14), h16(e12))) & (h16(op1(e14, e11)) = op2(h16(e14), h16(e11))) & (h16(op1(e14, e10)) = op2(h16(e14), h16(e10))) & (h16(op1(e13, e14)) = op2(h16(e13), h16(e14))) & (h16(op1(e13, e13)) = op2(h16(e13), h16(e13))) & (h16(op1(e13, e12)) = op2(h16(e13), h16(e12))) & (h16(op1(e13, e11)) = op2(h16(e13), h16(e11))) & (h16(op1(e13, e10)) = op2(h16(e13), h16(e10))) & (h16(op1(e12, e14)) = op2(h16(e12), h16(e14))) & (h16(op1(e12, e13)) = op2(h16(e12), h16(e13))) & (h16(op1(e12, e12)) = op2(h16(e12), h16(e12))) & (h16(op1(e12, e11)) = op2(h16(e12), h16(e11))) & (h16(op1(e12, e10)) = op2(h16(e12), h16(e10))) & (h16(op1(e11, e14)) = op2(h16(e11), h16(e14))) & (h16(op1(e11, e13)) = op2(h16(e11), h16(e13))) & (h16(op1(e11, e12)) = op2(h16(e11), h16(e12))) & (h16(op1(e11, e11)) = op2(h16(e11), h16(e11))) & (h16(op1(e11, e10)) = op2(h16(e11), h16(e10))) & (h16(op1(e10, e14)) = op2(h16(e10), h16(e14))) & (h16(op1(e10, e13)) = op2(h16(e10), h16(e13))) & (h16(op1(e10, e12)) = op2(h16(e10), h16(e12))) & (h16(op1(e10, e11)) = op2(h16(e10), h16(e11))) & (h16(op1(e10, e10)) = op2(h16(e10), h16(e10)))) | (((e24 = h15(e14)) | (e24 = h15(e13)) | (e24 = h15(e12)) | (e24 = h15(e11)) | (e24 = h15(e10))) & ((e23 = h15(e14)) | (e23 = h15(e13)) | (e23 = h15(e12)) | (e23 = h15(e11)) | (e23 = h15(e10))) & ((e22 = h15(e14)) | (e22 = h15(e13)) | (e22 = h15(e12)) | (e22 = h15(e11)) | (e22 = h15(e10))) & ((e21 = h15(e14)) | (e21 = h15(e13)) | (e21 = h15(e12)) | (e21 = h15(e11)) | (e21 = h15(e10))) & ((e20 = h15(e14)) | (e20 = h15(e13)) | (e20 = h15(e12)) | (e20 = h15(e11)) | (e20 = h15(e10))) & (h15(op1(e14, e14)) = op2(h15(e14), h15(e14))) & (h15(op1(e14, e13)) = op2(h15(e14), h15(e13))) & (h15(op1(e14, e12)) = op2(h15(e14), h15(e12))) & (h15(op1(e14, e11)) = op2(h15(e14), h15(e11))) & (h15(op1(e14, e10)) = op2(h15(e14), h15(e10))) & (h15(op1(e13, e14)) = op2(h15(e13), h15(e14))) & (h15(op1(e13, e13)) = op2(h15(e13), h15(e13))) & (h15(op1(e13, e12)) = op2(h15(e13), h15(e12))) & (h15(op1(e13, e11)) = op2(h15(e13), h15(e11))) & (h15(op1(e13, e10)) = op2(h15(e13), h15(e10))) & (h15(op1(e12, e14)) = op2(h15(e12), h15(e14))) & (h15(op1(e12, e13)) = op2(h15(e12), h15(e13))) & (h15(op1(e12, e12)) = op2(h15(e12), h15(e12))) & (h15(op1(e12, e11)) = op2(h15(e12), h15(e11))) & (h15(op1(e12, e10)) = op2(h15(e12), h15(e10))) & (h15(op1(e11, e14)) = op2(h15(e11), h15(e14))) & (h15(op1(e11, e13)) = op2(h15(e11), h15(e13))) & (h15(op1(e11, e12)) = op2(h15(e11), h15(e12))) & (h15(op1(e11, e11)) = op2(h15(e11), h15(e11))) & (h15(op1(e11, e10)) = op2(h15(e11), h15(e10))) & (h15(op1(e10, e14)) = op2(h15(e10), h15(e14))) & (h15(op1(e10, e13)) = op2(h15(e10), h15(e13))) & (h15(op1(e10, e12)) = op2(h15(e10), h15(e12))) & (h15(op1(e10, e11)) = op2(h15(e10), h15(e11))) & (h15(op1(e10, e10)) = op2(h15(e10), h15(e10)))) | (((e24 = h14(e14)) | (e24 = h14(e13)) | (e24 = h14(e12)) | (e24 = h14(e11)) | (e24 = h14(e10))) & ((e23 = h14(e14)) | (e23 = h14(e13)) | (e23 = h14(e12)) | (e23 = h14(e11)) | (e23 = h14(e10))) & ((e22 = h14(e14)) | (e22 = h14(e13)) | (e22 = h14(e12)) | (e22 = h14(e11)) | (e22 = h14(e10))) & ((e21 = h14(e14)) | (e21 = h14(e13)) | (e21 = h14(e12)) | (e21 = h14(e11)) | (e21 = h14(e10))) & ((e20 = h14(e14)) | (e20 = h14(e13)) | (e20 = h14(e12)) | (e20 = h14(e11)) | (e20 = h14(e10))) & (h14(op1(e14, e14)) = op2(h14(e14), h14(e14))) & (h14(op1(e14, e13)) = op2(h14(e14), h14(e13))) & (h14(op1(e14, e12)) = op2(h14(e14), h14(e12))) & (h14(op1(e14, e11)) = op2(h14(e14), h14(e11))) & (h14(op1(e14, e10)) = op2(h14(e14), h14(e10))) & (h14(op1(e13, e14)) = op2(h14(e13), h14(e14))) & (h14(op1(e13, e13)) = op2(h14(e13), h14(e13))) & (h14(op1(e13, e12)) = op2(h14(e13), h14(e12))) & (h14(op1(e13, e11)) = op2(h14(e13), h14(e11))) & (h14(op1(e13, e10)) = op2(h14(e13), h14(e10))) & (h14(op1(e12, e14)) = op2(h14(e12), h14(e14))) & (h14(op1(e12, e13)) = op2(h14(e12), h14(e13))) & (h14(op1(e12, e12)) = op2(h14(e12), h14(e12))) & (h14(op1(e12, e11)) = op2(h14(e12), h14(e11))) & (h14(op1(e12, e10)) = op2(h14(e12), h14(e10))) & (h14(op1(e11, e14)) = op2(h14(e11), h14(e14))) & (h14(op1(e11, e13)) = op2(h14(e11), h14(e13))) & (h14(op1(e11, e12)) = op2(h14(e11), h14(e12))) & (h14(op1(e11, e11)) = op2(h14(e11), h14(e11))) & (h14(op1(e11, e10)) = op2(h14(e11), h14(e10))) & (h14(op1(e10, e14)) = op2(h14(e10), h14(e14))) & (h14(op1(e10, e13)) = op2(h14(e10), h14(e13))) & (h14(op1(e10, e12)) = op2(h14(e10), h14(e12))) & (h14(op1(e10, e11)) = op2(h14(e10), h14(e11))) & (h14(op1(e10, e10)) = op2(h14(e10), h14(e10)))) | (((e24 = h13(e14)) | (e24 = h13(e13)) | (e24 = h13(e12)) | (e24 = h13(e11)) | (e24 = h13(e10))) & ((e23 = h13(e14)) | (e23 = h13(e13)) | (e23 = h13(e12)) | (e23 = h13(e11)) | (e23 = h13(e10))) & ((e22 = h13(e14)) | (e22 = h13(e13)) | (e22 = h13(e12)) | (e22 = h13(e11)) | (e22 = h13(e10))) & ((e21 = h13(e14)) | (e21 = h13(e13)) | (e21 = h13(e12)) | (e21 = h13(e11)) | (e21 = h13(e10))) & ((e20 = h13(e14)) | (e20 = h13(e13)) | (e20 = h13(e12)) | (e20 = h13(e11)) | (e20 = h13(e10))) & (h13(op1(e14, e14)) = op2(h13(e14), h13(e14))) & (h13(op1(e14, e13)) = op2(h13(e14), h13(e13))) & (h13(op1(e14, e12)) = op2(h13(e14), h13(e12))) & (h13(op1(e14, e11)) = op2(h13(e14), h13(e11))) & (h13(op1(e14, e10)) = op2(h13(e14), h13(e10))) & (h13(op1(e13, e14)) = op2(h13(e13), h13(e14))) & (h13(op1(e13, e13)) = op2(h13(e13), h13(e13))) & (h13(op1(e13, e12)) = op2(h13(e13), h13(e12))) & (h13(op1(e13, e11)) = op2(h13(e13), h13(e11))) & (h13(op1(e13, e10)) = op2(h13(e13), h13(e10))) & (h13(op1(e12, e14)) = op2(h13(e12), h13(e14))) & (h13(op1(e12, e13)) = op2(h13(e12), h13(e13))) & (h13(op1(e12, e12)) = op2(h13(e12), h13(e12))) & (h13(op1(e12, e11)) = op2(h13(e12), h13(e11))) & (h13(op1(e12, e10)) = op2(h13(e12), h13(e10))) & (h13(op1(e11, e14)) = op2(h13(e11), h13(e14))) & (h13(op1(e11, e13)) = op2(h13(e11), h13(e13))) & (h13(op1(e11, e12)) = op2(h13(e11), h13(e12))) & (h13(op1(e11, e11)) = op2(h13(e11), h13(e11))) & (h13(op1(e11, e10)) = op2(h13(e11), h13(e10))) & (h13(op1(e10, e14)) = op2(h13(e10), h13(e14))) & (h13(op1(e10, e13)) = op2(h13(e10), h13(e13))) & (h13(op1(e10, e12)) = op2(h13(e10), h13(e12))) & (h13(op1(e10, e11)) = op2(h13(e10), h13(e11))) & (h13(op1(e10, e10)) = op2(h13(e10), h13(e10)))) | (((e24 = h12(e14)) | (e24 = h12(e13)) | (e24 = h12(e12)) | (e24 = h12(e11)) | (e24 = h12(e10))) & ((e23 = h12(e14)) | (e23 = h12(e13)) | (e23 = h12(e12)) | (e23 = h12(e11)) | (e23 = h12(e10))) & ((e22 = h12(e14)) | (e22 = h12(e13)) | (e22 = h12(e12)) | (e22 = h12(e11)) | (e22 = h12(e10))) & ((e21 = h12(e14)) | (e21 = h12(e13)) | (e21 = h12(e12)) | (e21 = h12(e11)) | (e21 = h12(e10))) & ((e20 = h12(e14)) | (e20 = h12(e13)) | (e20 = h12(e12)) | (e20 = h12(e11)) | (e20 = h12(e10))) & (h12(op1(e14, e14)) = op2(h12(e14), h12(e14))) & (h12(op1(e14, e13)) = op2(h12(e14), h12(e13))) & (h12(op1(e14, e12)) = op2(h12(e14), h12(e12))) & (h12(op1(e14, e11)) = op2(h12(e14), h12(e11))) & (h12(op1(e14, e10)) = op2(h12(e14), h12(e10))) & (h12(op1(e13, e14)) = op2(h12(e13), h12(e14))) & (h12(op1(e13, e13)) = op2(h12(e13), h12(e13))) & (h12(op1(e13, e12)) = op2(h12(e13), h12(e12))) & (h12(op1(e13, e11)) = op2(h12(e13), h12(e11))) & (h12(op1(e13, e10)) = op2(h12(e13), h12(e10))) & (h12(op1(e12, e14)) = op2(h12(e12), h12(e14))) & (h12(op1(e12, e13)) = op2(h12(e12), h12(e13))) & (h12(op1(e12, e12)) = op2(h12(e12), h12(e12))) & (h12(op1(e12, e11)) = op2(h12(e12), h12(e11))) & (h12(op1(e12, e10)) = op2(h12(e12), h12(e10))) & (h12(op1(e11, e14)) = op2(h12(e11), h12(e14))) & (h12(op1(e11, e13)) = op2(h12(e11), h12(e13))) & (h12(op1(e11, e12)) = op2(h12(e11), h12(e12))) & (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) & (h12(op1(e11, e10)) = op2(h12(e11), h12(e10))) & (h12(op1(e10, e14)) = op2(h12(e10), h12(e14))) & (h12(op1(e10, e13)) = op2(h12(e10), h12(e13))) & (h12(op1(e10, e12)) = op2(h12(e10), h12(e12))) & (h12(op1(e10, e11)) = op2(h12(e10), h12(e11))) & (h12(op1(e10, e10)) = op2(h12(e10), h12(e10)))) | (((e24 = h11(e14)) | (e24 = h11(e13)) | (e24 = h11(e12)) | (e24 = h11(e11)) | (e24 = h11(e10))) & ((e23 = h11(e14)) | (e23 = h11(e13)) | (e23 = h11(e12)) | (e23 = h11(e11)) | (e23 = h11(e10))) & ((e22 = h11(e14)) | (e22 = h11(e13)) | (e22 = h11(e12)) | (e22 = h11(e11)) | (e22 = h11(e10))) & ((e21 = h11(e14)) | (e21 = h11(e13)) | (e21 = h11(e12)) | (e21 = h11(e11)) | (e21 = h11(e10))) & ((e20 = h11(e14)) | (e20 = h11(e13)) | (e20 = h11(e12)) | (e20 = h11(e11)) | (e20 = h11(e10))) & (h11(op1(e14, e14)) = op2(h11(e14), h11(e14))) & (h11(op1(e14, e13)) = op2(h11(e14), h11(e13))) & (h11(op1(e14, e12)) = op2(h11(e14), h11(e12))) & (h11(op1(e14, e11)) = op2(h11(e14), h11(e11))) & (h11(op1(e14, e10)) = op2(h11(e14), h11(e10))) & (h11(op1(e13, e14)) = op2(h11(e13), h11(e14))) & (h11(op1(e13, e13)) = op2(h11(e13), h11(e13))) & (h11(op1(e13, e12)) = op2(h11(e13), h11(e12))) & (h11(op1(e13, e11)) = op2(h11(e13), h11(e11))) & (h11(op1(e13, e10)) = op2(h11(e13), h11(e10))) & (h11(op1(e12, e14)) = op2(h11(e12), h11(e14))) & (h11(op1(e12, e13)) = op2(h11(e12), h11(e13))) & (h11(op1(e12, e12)) = op2(h11(e12), h11(e12))) & (h11(op1(e12, e11)) = op2(h11(e12), h11(e11))) & (h11(op1(e12, e10)) = op2(h11(e12), h11(e10))) & (h11(op1(e11, e14)) = op2(h11(e11), h11(e14))) & (h11(op1(e11, e13)) = op2(h11(e11), h11(e13))) & (h11(op1(e11, e12)) = op2(h11(e11), h11(e12))) & (h11(op1(e11, e11)) = op2(h11(e11), h11(e11))) & (h11(op1(e11, e10)) = op2(h11(e11), h11(e10))) & (h11(op1(e10, e14)) = op2(h11(e10), h11(e14))) & (h11(op1(e10, e13)) = op2(h11(e10), h11(e13))) & (h11(op1(e10, e12)) = op2(h11(e10), h11(e12))) & (h11(op1(e10, e11)) = op2(h11(e10), h11(e11))) & (h11(op1(e10, e10)) = op2(h11(e10), h11(e10)))) | (((e24 = h10(e14)) | (e24 = h10(e13)) | (e24 = h10(e12)) | (e24 = h10(e11)) | (e24 = h10(e10))) & ((e23 = h10(e14)) | (e23 = h10(e13)) | (e23 = h10(e12)) | (e23 = h10(e11)) | (e23 = h10(e10))) & ((e22 = h10(e14)) | (e22 = h10(e13)) | (e22 = h10(e12)) | (e22 = h10(e11)) | (e22 = h10(e10))) & ((e21 = h10(e14)) | (e21 = h10(e13)) | (e21 = h10(e12)) | (e21 = h10(e11)) | (e21 = h10(e10))) & ((e20 = h10(e14)) | (e20 = h10(e13)) | (e20 = h10(e12)) | (e20 = h10(e11)) | (e20 = h10(e10))) & (h10(op1(e14, e14)) = op2(h10(e14), h10(e14))) & (h10(op1(e14, e13)) = op2(h10(e14), h10(e13))) & (h10(op1(e14, e12)) = op2(h10(e14), h10(e12))) & (h10(op1(e14, e11)) = op2(h10(e14), h10(e11))) & (h10(op1(e14, e10)) = op2(h10(e14), h10(e10))) & (h10(op1(e13, e14)) = op2(h10(e13), h10(e14))) & (h10(op1(e13, e13)) = op2(h10(e13), h10(e13))) & (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) & (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) & (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) & (h10(op1(e12, e14)) = op2(h10(e12), h10(e14))) & (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) & (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) & (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) & (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) & (h10(op1(e11, e14)) = op2(h10(e11), h10(e14))) & (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) & (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) & (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) & (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) & (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) & (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) & (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) & (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) & (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))) | (((e24 = h9(e14)) | (e24 = h9(e13)) | (e24 = h9(e12)) | (e24 = h9(e11)) | (e24 = h9(e10))) & ((e23 = h9(e14)) | (e23 = h9(e13)) | (e23 = h9(e12)) | (e23 = h9(e11)) | (e23 = h9(e10))) & ((e22 = h9(e14)) | (e22 = h9(e13)) | (e22 = h9(e12)) | (e22 = h9(e11)) | (e22 = h9(e10))) & ((e21 = h9(e14)) | (e21 = h9(e13)) | (e21 = h9(e12)) | (e21 = h9(e11)) | (e21 = h9(e10))) & ((e20 = h9(e14)) | (e20 = h9(e13)) | (e20 = h9(e12)) | (e20 = h9(e11)) | (e20 = h9(e10))) & (h9(op1(e14, e14)) = op2(h9(e14), h9(e14))) & (h9(op1(e14, e13)) = op2(h9(e14), h9(e13))) & (h9(op1(e14, e12)) = op2(h9(e14), h9(e12))) & (h9(op1(e14, e11)) = op2(h9(e14), h9(e11))) & (h9(op1(e14, e10)) = op2(h9(e14), h9(e10))) & (h9(op1(e13, e14)) = op2(h9(e13), h9(e14))) & (h9(op1(e13, e13)) = op2(h9(e13), h9(e13))) & (h9(op1(e13, e12)) = op2(h9(e13), h9(e12))) & (h9(op1(e13, e11)) = op2(h9(e13), h9(e11))) & (h9(op1(e13, e10)) = op2(h9(e13), h9(e10))) & (h9(op1(e12, e14)) = op2(h9(e12), h9(e14))) & (h9(op1(e12, e13)) = op2(h9(e12), h9(e13))) & (h9(op1(e12, e12)) = op2(h9(e12), h9(e12))) & (h9(op1(e12, e11)) = op2(h9(e12), h9(e11))) & (h9(op1(e12, e10)) = op2(h9(e12), h9(e10))) & (h9(op1(e11, e14)) = op2(h9(e11), h9(e14))) & (h9(op1(e11, e13)) = op2(h9(e11), h9(e13))) & (h9(op1(e11, e12)) = op2(h9(e11), h9(e12))) & (h9(op1(e11, e11)) = op2(h9(e11), h9(e11))) & (h9(op1(e11, e10)) = op2(h9(e11), h9(e10))) & (h9(op1(e10, e14)) = op2(h9(e10), h9(e14))) & (h9(op1(e10, e13)) = op2(h9(e10), h9(e13))) & (h9(op1(e10, e12)) = op2(h9(e10), h9(e12))) & (h9(op1(e10, e11)) = op2(h9(e10), h9(e11))) & (h9(op1(e10, e10)) = op2(h9(e10), h9(e10)))) | (((e24 = h8(e14)) | (e24 = h8(e13)) | (e24 = h8(e12)) | (e24 = h8(e11)) | (e24 = h8(e10))) & ((e23 = h8(e14)) | (e23 = h8(e13)) | (e23 = h8(e12)) | (e23 = h8(e11)) | (e23 = h8(e10))) & ((e22 = h8(e14)) | (e22 = h8(e13)) | (e22 = h8(e12)) | (e22 = h8(e11)) | (e22 = h8(e10))) & ((e21 = h8(e14)) | (e21 = h8(e13)) | (e21 = h8(e12)) | (e21 = h8(e11)) | (e21 = h8(e10))) & ((e20 = h8(e14)) | (e20 = h8(e13)) | (e20 = h8(e12)) | (e20 = h8(e11)) | (e20 = h8(e10))) & (h8(op1(e14, e14)) = op2(h8(e14), h8(e14))) & (h8(op1(e14, e13)) = op2(h8(e14), h8(e13))) & (h8(op1(e14, e12)) = op2(h8(e14), h8(e12))) & (h8(op1(e14, e11)) = op2(h8(e14), h8(e11))) & (h8(op1(e14, e10)) = op2(h8(e14), h8(e10))) & (h8(op1(e13, e14)) = op2(h8(e13), h8(e14))) & (h8(op1(e13, e13)) = op2(h8(e13), h8(e13))) & (h8(op1(e13, e12)) = op2(h8(e13), h8(e12))) & (h8(op1(e13, e11)) = op2(h8(e13), h8(e11))) & (h8(op1(e13, e10)) = op2(h8(e13), h8(e10))) & (h8(op1(e12, e14)) = op2(h8(e12), h8(e14))) & (h8(op1(e12, e13)) = op2(h8(e12), h8(e13))) & (h8(op1(e12, e12)) = op2(h8(e12), h8(e12))) & (h8(op1(e12, e11)) = op2(h8(e12), h8(e11))) & (h8(op1(e12, e10)) = op2(h8(e12), h8(e10))) & (h8(op1(e11, e14)) = op2(h8(e11), h8(e14))) & (h8(op1(e11, e13)) = op2(h8(e11), h8(e13))) & (h8(op1(e11, e12)) = op2(h8(e11), h8(e12))) & (h8(op1(e11, e11)) = op2(h8(e11), h8(e11))) & (h8(op1(e11, e10)) = op2(h8(e11), h8(e10))) & (h8(op1(e10, e14)) = op2(h8(e10), h8(e14))) & (h8(op1(e10, e13)) = op2(h8(e10), h8(e13))) & (h8(op1(e10, e12)) = op2(h8(e10), h8(e12))) & (h8(op1(e10, e11)) = op2(h8(e10), h8(e11))) & (h8(op1(e10, e10)) = op2(h8(e10), h8(e10)))) | (((e24 = h7(e14)) | (e24 = h7(e13)) | (e24 = h7(e12)) | (e24 = h7(e11)) | (e24 = h7(e10))) & ((e23 = h7(e14)) | (e23 = h7(e13)) | (e23 = h7(e12)) | (e23 = h7(e11)) | (e23 = h7(e10))) & ((e22 = h7(e14)) | (e22 = h7(e13)) | (e22 = h7(e12)) | (e22 = h7(e11)) | (e22 = h7(e10))) & ((e21 = h7(e14)) | (e21 = h7(e13)) | (e21 = h7(e12)) | (e21 = h7(e11)) | (e21 = h7(e10))) & ((e20 = h7(e14)) | (e20 = h7(e13)) | (e20 = h7(e12)) | (e20 = h7(e11)) | (e20 = h7(e10))) & (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) & (h7(op1(e14, e13)) = op2(h7(e14), h7(e13))) & (h7(op1(e14, e12)) = op2(h7(e14), h7(e12))) & (h7(op1(e14, e11)) = op2(h7(e14), h7(e11))) & (h7(op1(e14, e10)) = op2(h7(e14), h7(e10))) & (h7(op1(e13, e14)) = op2(h7(e13), h7(e14))) & (h7(op1(e13, e13)) = op2(h7(e13), h7(e13))) & (h7(op1(e13, e12)) = op2(h7(e13), h7(e12))) & (h7(op1(e13, e11)) = op2(h7(e13), h7(e11))) & (h7(op1(e13, e10)) = op2(h7(e13), h7(e10))) & (h7(op1(e12, e14)) = op2(h7(e12), h7(e14))) & (h7(op1(e12, e13)) = op2(h7(e12), h7(e13))) & (h7(op1(e12, e12)) = op2(h7(e12), h7(e12))) & (h7(op1(e12, e11)) = op2(h7(e12), h7(e11))) & (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) & (h7(op1(e11, e14)) = op2(h7(e11), h7(e14))) & (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) & (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) & (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) & (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) & (h7(op1(e10, e14)) = op2(h7(e10), h7(e14))) & (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) & (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) & (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) & (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))) | (((e24 = h6(e14)) | (e24 = h6(e13)) | (e24 = h6(e12)) | (e24 = h6(e11)) | (e24 = h6(e10))) & ((e23 = h6(e14)) | (e23 = h6(e13)) | (e23 = h6(e12)) | (e23 = h6(e11)) | (e23 = h6(e10))) & ((e22 = h6(e14)) | (e22 = h6(e13)) | (e22 = h6(e12)) | (e22 = h6(e11)) | (e22 = h6(e10))) & ((e21 = h6(e14)) | (e21 = h6(e13)) | (e21 = h6(e12)) | (e21 = h6(e11)) | (e21 = h6(e10))) & ((e20 = h6(e14)) | (e20 = h6(e13)) | (e20 = h6(e12)) | (e20 = h6(e11)) | (e20 = h6(e10))) & (h6(op1(e14, e14)) = op2(h6(e14), h6(e14))) & (h6(op1(e14, e13)) = op2(h6(e14), h6(e13))) & (h6(op1(e14, e12)) = op2(h6(e14), h6(e12))) & (h6(op1(e14, e11)) = op2(h6(e14), h6(e11))) & (h6(op1(e14, e10)) = op2(h6(e14), h6(e10))) & (h6(op1(e13, e14)) = op2(h6(e13), h6(e14))) & (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) & (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) & (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) & (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) & (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) & (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) & (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) & (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) & (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) & (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) & (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) & (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) & (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) & (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) & (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) & (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) & (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) & (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) & (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))) | (((e24 = h5(e14)) | (e24 = h5(e13)) | (e24 = h5(e12)) | (e24 = h5(e11)) | (e24 = h5(e10))) & ((e23 = h5(e14)) | (e23 = h5(e13)) | (e23 = h5(e12)) | (e23 = h5(e11)) | (e23 = h5(e10))) & ((e22 = h5(e14)) | (e22 = h5(e13)) | (e22 = h5(e12)) | (e22 = h5(e11)) | (e22 = h5(e10))) & ((e21 = h5(e14)) | (e21 = h5(e13)) | (e21 = h5(e12)) | (e21 = h5(e11)) | (e21 = h5(e10))) & ((e20 = h5(e14)) | (e20 = h5(e13)) | (e20 = h5(e12)) | (e20 = h5(e11)) | (e20 = h5(e10))) & (h5(op1(e14, e14)) = op2(h5(e14), h5(e14))) & (h5(op1(e14, e13)) = op2(h5(e14), h5(e13))) & (h5(op1(e14, e12)) = op2(h5(e14), h5(e12))) & (h5(op1(e14, e11)) = op2(h5(e14), h5(e11))) & (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) & (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) & (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) & (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) & (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) & (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) & (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) & (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) & (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) & (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) & (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) & (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) & (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) & (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) & (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) & (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) & (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) & (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) & (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) & (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) & (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))) | (((e24 = h4(e14)) | (e24 = h4(e13)) | (e24 = h4(e12)) | (e24 = h4(e11)) | (e24 = h4(e10))) & ((e23 = h4(e14)) | (e23 = h4(e13)) | (e23 = h4(e12)) | (e23 = h4(e11)) | (e23 = h4(e10))) & ((e22 = h4(e14)) | (e22 = h4(e13)) | (e22 = h4(e12)) | (e22 = h4(e11)) | (e22 = h4(e10))) & ((e21 = h4(e14)) | (e21 = h4(e13)) | (e21 = h4(e12)) | (e21 = h4(e11)) | (e21 = h4(e10))) & ((e20 = h4(e14)) | (e20 = h4(e13)) | (e20 = h4(e12)) | (e20 = h4(e11)) | (e20 = h4(e10))) & (h4(op1(e14, e14)) = op2(h4(e14), h4(e14))) & (h4(op1(e14, e13)) = op2(h4(e14), h4(e13))) & (h4(op1(e14, e12)) = op2(h4(e14), h4(e12))) & (h4(op1(e14, e11)) = op2(h4(e14), h4(e11))) & (h4(op1(e14, e10)) = op2(h4(e14), h4(e10))) & (h4(op1(e13, e14)) = op2(h4(e13), h4(e14))) & (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) & (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) & (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) & (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) & (h4(op1(e12, e14)) = op2(h4(e12), h4(e14))) & (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) & (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) & (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) & (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) & (h4(op1(e11, e14)) = op2(h4(e11), h4(e14))) & (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) & (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) & (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) & (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) & (h4(op1(e10, e14)) = op2(h4(e10), h4(e14))) & (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) & (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) & (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) & (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) | (((e24 = h3(e14)) | (e24 = h3(e13)) | (e24 = h3(e12)) | (e24 = h3(e11)) | (e24 = h3(e10))) & ((e23 = h3(e14)) | (e23 = h3(e13)) | (e23 = h3(e12)) | (e23 = h3(e11)) | (e23 = h3(e10))) & ((e22 = h3(e14)) | (e22 = h3(e13)) | (e22 = h3(e12)) | (e22 = h3(e11)) | (e22 = h3(e10))) & ((e21 = h3(e14)) | (e21 = h3(e13)) | (e21 = h3(e12)) | (e21 = h3(e11)) | (e21 = h3(e10))) & ((e20 = h3(e14)) | (e20 = h3(e13)) | (e20 = h3(e12)) | (e20 = h3(e11)) | (e20 = h3(e10))) & (h3(op1(e14, e14)) = op2(h3(e14), h3(e14))) & (h3(op1(e14, e13)) = op2(h3(e14), h3(e13))) & (h3(op1(e14, e12)) = op2(h3(e14), h3(e12))) & (h3(op1(e14, e11)) = op2(h3(e14), h3(e11))) & (h3(op1(e14, e10)) = op2(h3(e14), h3(e10))) & (h3(op1(e13, e14)) = op2(h3(e13), h3(e14))) & (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) & (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) & (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) & (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) & (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) & (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) & (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) & (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) & (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) & (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) & (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) & (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) & (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) & (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) & (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) & (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) & (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) & (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) & (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) | (((e24 = h2(e14)) | (e24 = h2(e13)) | (e24 = h2(e12)) | (e24 = h2(e11)) | (e24 = h2(e10))) & ((e23 = h2(e14)) | (e23 = h2(e13)) | (e23 = h2(e12)) | (e23 = h2(e11)) | (e23 = h2(e10))) & ((e22 = h2(e14)) | (e22 = h2(e13)) | (e22 = h2(e12)) | (e22 = h2(e11)) | (e22 = h2(e10))) & ((e21 = h2(e14)) | (e21 = h2(e13)) | (e21 = h2(e12)) | (e21 = h2(e11)) | (e21 = h2(e10))) & ((e20 = h2(e14)) | (e20 = h2(e13)) | (e20 = h2(e12)) | (e20 = h2(e11)) | (e20 = h2(e10))) & (h2(op1(e14, e14)) = op2(h2(e14), h2(e14))) & (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) & (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) & (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) & (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) & (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) & (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) & (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) & (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) & (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) & (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) & (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) & (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) & (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) & (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) & (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) & (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) & (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) & (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) & (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) & (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) & (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) & (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) & (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) & (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) | (((e24 = h1(e14)) | (e24 = h1(e13)) | (e24 = h1(e12)) | (e24 = h1(e11)) | (e24 = h1(e10))) & ((e23 = h1(e14)) | (e23 = h1(e13)) | (e23 = h1(e12)) | (e23 = h1(e11)) | (e23 = h1(e10))) & ((e22 = h1(e14)) | (e22 = h1(e13)) | (e22 = h1(e12)) | (e22 = h1(e11)) | (e22 = h1(e10))) & ((e21 = h1(e14)) | (e21 = h1(e13)) | (e21 = h1(e12)) | (e21 = h1(e11)) | (e21 = h1(e10))) & ((e20 = h1(e14)) | (e20 = h1(e13)) | (e20 = h1(e12)) | (e20 = h1(e11)) | (e20 = h1(e10))) & (h1(op1(e14, e14)) = op2(h1(e14), h1(e14))) & (h1(op1(e14, e13)) = op2(h1(e14), h1(e13))) & (h1(op1(e14, e12)) = op2(h1(e14), h1(e12))) & (h1(op1(e14, e11)) = op2(h1(e14), h1(e11))) & (h1(op1(e14, e10)) = op2(h1(e14), h1(e10))) & (h1(op1(e13, e14)) = op2(h1(e13), h1(e14))) & (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) & (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) & (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) & (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) & (h1(op1(e12, e14)) = op2(h1(e12), h1(e14))) & (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) & (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) & (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) & (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) & (h1(op1(e11, e14)) = op2(h1(e11), h1(e14))) & (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) & (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) & (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) & (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) & (h1(op1(e10, e14)) = op2(h1(e10), h1(e14))) & (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) & (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) & (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) & (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG049+1.p', co1)).
fof(f2860, plain, (~ spl80_293 | ~ spl80_296), inference(avatar_split_clause, [], [f828, f2857, f2841])).
fof(f828, plain, (~ (e20 = h18(e10)) | ~ sP68), inference(cnf_transformation, [], [f131])).
fof(f131, plain, ((~ (e20 = h18(e14)) & ~ (e20 = h18(e13)) & ~ (e20 = h18(e12)) & ~ (e20 = h18(e11)) & ~ (e20 = h18(e10))) | ~ sP68), inference(nnf_transformation, [], [f107])).
fof(f2831, plain, (~ spl80_289 | ~ spl80_291), inference(avatar_split_clause, [], [f2826, f2828, f2817])).
fof(f2826, plain, (~ (e22 = h8(e13)) | ~ sP70), inference(forward_demodulation, [], [f819, f2643])).
fof(f819, plain, (~ (e22 = h18(e11)) | ~ sP70), inference(cnf_transformation, [], [f129])).
fof(f129, plain, ((~ (e22 = h18(e14)) & ~ (e22 = h18(e13)) & ~ (e22 = h18(e12)) & ~ (e22 = h18(e11)) & ~ (e22 = h18(e10))) | ~ sP70), inference(nnf_transformation, [], [f109])).
fof(f2801, plain, (~ spl80_285 | ~ spl80_286), inference(avatar_split_clause, [], [f2792, f2798, f2794])).
fof(f2792, plain, (~ (e23 = h8(e11)) | ~ sP71), inference(forward_demodulation, [], [f816, f2642])).
fof(f816, plain, (~ (e23 = h18(e13)) | ~ sP71), inference(cnf_transformation, [], [f128])).
fof(f128, plain, ((~ (e23 = h18(e14)) & ~ (e23 = h18(e13)) & ~ (e23 = h18(e12)) & ~ (e23 = h18(e11)) & ~ (e23 = h18(e10))) | ~ sP71), inference(nnf_transformation, [], [f110])).
fof(f2508, plain, spl80_142, inference(avatar_split_clause, [], [f671, f1916])).
fof(f671, plain, (e21 = op2(e24, e22)), inference(cnf_transformation, [], [f15])).
fof(f2507, plain, spl80_184, inference(avatar_split_clause, [], [f672, f2092])).
fof(f672, plain, (e23 = op2(e22, e24)), inference(cnf_transformation, [], [f15])).
fof(f2506, plain, spl80_12, inference(avatar_split_clause, [], [f668, f1320])).
fof(f668, plain, (e11 = op1(e14, e12)), inference(cnf_transformation, [], [f14])).
fof(f2505, plain, spl80_54, inference(avatar_split_clause, [], [f669, f1496])).
fof(f669, plain, (e13 = op1(e12, e14)), inference(cnf_transformation, [], [f14])).
fof(f2500, plain, (~ spl80_221 | spl80_227), inference(avatar_split_clause, [], [f647, f2273, f2248])).
fof(f647, plain, ((e21 = op2(e21, e20)) | ~ (e20 = op2(e21, e21))), inference(cnf_transformation, [], [f13])).
fof(f13, plain, (((e24 = op2(e24, e24)) | ~ (e24 = op2(e24, e24))) & ((e24 = op2(e24, e23)) | ~ (e23 = op2(e24, e24))) & ((e24 = op2(e24, e22)) | ~ (e22 = op2(e24, e24))) & ((e24 = op2(e24, e21)) | ~ (e21 = op2(e24, e24))) & ((e24 = op2(e24, e20)) | ~ (e20 = op2(e24, e24))) & ((e23 = op2(e23, e24)) | ~ (e24 = op2(e23, e23))) & ((e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23))) & ((e23 = op2(e23, e22)) | ~ (e22 = op2(e23, e23))) & ((e23 = op2(e23, e21)) | ~ (e21 = op2(e23, e23))) & ((e23 = op2(e23, e20)) | ~ (e20 = op2(e23, e23))) & ((e22 = op2(e22, e24)) | ~ (e24 = op2(e22, e22))) & ((e22 = op2(e22, e23)) | ~ (e23 = op2(e22, e22))) & ((e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & ((e22 = op2(e22, e21)) | ~ (e21 = op2(e22, e22))) & ((e22 = op2(e22, e20)) | ~ (e20 = op2(e22, e22))) & ((e21 = op2(e21, e24)) | ~ (e24 = op2(e21, e21))) & ((e21 = op2(e21, e23)) | ~ (e23 = op2(e21, e21))) & ((e21 = op2(e21, e22)) | ~ (e22 = op2(e21, e21))) & ((e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & ((e21 = op2(e21, e20)) | ~ (e20 = op2(e21, e21))) & ((e20 = op2(e20, e24)) | ~ (op2(e20, e20) = e24)) & ((e20 = op2(e20, e23)) | ~ (op2(e20, e20) = e23)) & ((e20 = op2(e20, e22)) | ~ (op2(e20, e20) = e22)) & ((e20 = op2(e20, e21)) | ~ (op2(e20, e20) = e21)) & ((e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG049+1.p', ax13)).
fof(f2499, plain, (~ spl80_223 | spl80_217), inference(avatar_split_clause, [], [f649, f2231, f2256])).
fof(f2256, plain, (spl80_223 <=> (e22 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl80_223])])).
fof(f649, plain, ((e21 = op2(e21, e22)) | ~ (e22 = op2(e21, e21))), inference(cnf_transformation, [], [f13])).
fof(f2498, plain, (~ spl80_224 | spl80_212), inference(avatar_split_clause, [], [f650, f2210, f2260])).
fof(f2260, plain, (spl80_224 <=> (e23 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl80_224])])).
fof(f650, plain, ((e21 = op2(e21, e23)) | ~ (e23 = op2(e21, e21))), inference(cnf_transformation, [], [f13])).
fof(f2497, plain, (~ spl80_225 | spl80_207), inference(avatar_split_clause, [], [f651, f2189, f2264])).
fof(f2264, plain, (spl80_225 <=> (e24 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl80_225])])).
fof(f651, plain, ((e21 = op2(e21, e24)) | ~ (e24 = op2(e21, e21))), inference(cnf_transformation, [], [f13])).
fof(f2493, plain, (~ spl80_195 | spl80_183), inference(avatar_split_clause, [], [f656, f2088, f2138])).
fof(f2138, plain, (spl80_195 <=> (e24 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl80_195])])).
fof(f656, plain, ((e22 = op2(e22, e24)) | ~ (e24 = op2(e22, e22))), inference(cnf_transformation, [], [f13])).
fof(f2491, plain, (~ spl80_162 | spl80_174), inference(avatar_split_clause, [], [f658, f2050, f2000])).
fof(f2000, plain, (spl80_162 <=> (e21 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl80_162])])).
fof(f658, plain, ((e23 = op2(e23, e21)) | ~ (e21 = op2(e23, e23))), inference(cnf_transformation, [], [f13])).
fof(f2490, plain, (~ spl80_163 | spl80_169), inference(avatar_split_clause, [], [f659, f2029, f2004])).
fof(f2004, plain, (spl80_163 <=> (e22 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl80_163])])).
fof(f659, plain, ((e23 = op2(e23, e22)) | ~ (e22 = op2(e23, e23))), inference(cnf_transformation, [], [f13])).
fof(f2480, plain, (~ spl80_91 | spl80_97), inference(avatar_split_clause, [], [f622, f1677, f1652])).
fof(f622, plain, ((e11 = op1(e11, e10)) | ~ (e10 = op1(e11, e11))), inference(cnf_transformation, [], [f12])).
fof(f12, plain, (((e14 = op1(e14, e14)) | ~ (e14 = op1(e14, e14))) & ((e14 = op1(e14, e13)) | ~ (e13 = op1(e14, e14))) & ((e14 = op1(e14, e12)) | ~ (e12 = op1(e14, e14))) & ((e14 = op1(e14, e11)) | ~ (e11 = op1(e14, e14))) & ((e14 = op1(e14, e10)) | ~ (e10 = op1(e14, e14))) & ((e13 = op1(e13, e14)) | ~ (e14 = op1(e13, e13))) & ((e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13))) & ((e13 = op1(e13, e12)) | ~ (e12 = op1(e13, e13))) & ((e13 = op1(e13, e11)) | ~ (e11 = op1(e13, e13))) & ((e13 = op1(e13, e10)) | ~ (e10 = op1(e13, e13))) & ((e12 = op1(e12, e14)) | ~ (e14 = op1(e12, e12))) & ((e12 = op1(e12, e13)) | ~ (e13 = op1(e12, e12))) & ((e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & ((e12 = op1(e12, e11)) | ~ (e11 = op1(e12, e12))) & ((e12 = op1(e12, e10)) | ~ (e10 = op1(e12, e12))) & ((e11 = op1(e11, e14)) | ~ (e14 = op1(e11, e11))) & ((e11 = op1(e11, e13)) | ~ (e13 = op1(e11, e11))) & ((e11 = op1(e11, e12)) | ~ (e12 = op1(e11, e11))) & ((e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & ((e11 = op1(e11, e10)) | ~ (e10 = op1(e11, e11))) & ((e10 = op1(e10, e14)) | ~ (op1(e10, e10) = e14)) & ((e10 = op1(e10, e13)) | ~ (op1(e10, e10) = e13)) & ((e10 = op1(e10, e12)) | ~ (op1(e10, e10) = e12)) & ((e10 = op1(e10, e11)) | ~ (op1(e10, e10) = e11)) & ((e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG049+1.p', ax12)).
fof(f2478, plain, (~ spl80_94 | spl80_82), inference(avatar_split_clause, [], [f625, f1614, f1664])).
fof(f1664, plain, (spl80_94 <=> (e13 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl80_94])])).
fof(f625, plain, ((e11 = op1(e11, e13)) | ~ (e13 = op1(e11, e11))), inference(cnf_transformation, [], [f12])).
fof(f2477, plain, (~ spl80_95 | spl80_77), inference(avatar_split_clause, [], [f626, f1593, f1668])).
fof(f1668, plain, (spl80_95 <=> (e14 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl80_95])])).
fof(f626, plain, ((e11 = op1(e11, e14)) | ~ (e14 = op1(e11, e11))), inference(cnf_transformation, [], [f12])).
fof(f2475, plain, (~ spl80_62 | spl80_68), inference(avatar_split_clause, [], [f628, f1555, f1530])).
fof(f1530, plain, (spl80_62 <=> (e11 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl80_62])])).
fof(f628, plain, ((e12 = op1(e12, e11)) | ~ (e11 = op1(e12, e12))), inference(cnf_transformation, [], [f12])).
fof(f2473, plain, (~ spl80_65 | spl80_53), inference(avatar_split_clause, [], [f631, f1492, f1542])).
fof(f1542, plain, (spl80_65 <=> (e14 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl80_65])])).
fof(f631, plain, ((e12 = op1(e12, e14)) | ~ (e14 = op1(e12, e12))), inference(cnf_transformation, [], [f12])).
fof(f2470, plain, (~ spl80_33 | spl80_39), inference(avatar_split_clause, [], [f634, f1433, f1408])).
fof(f1408, plain, (spl80_33 <=> (e12 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl80_33])])).
fof(f634, plain, ((e13 = op1(e13, e12)) | ~ (e12 = op1(e13, e13))), inference(cnf_transformation, [], [f12])).
fof(f2469, plain, (~ spl80_35 | spl80_29), inference(avatar_split_clause, [], [f636, f1391, f1416])).
fof(f1416, plain, (spl80_35 <=> (e14 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl80_35])])).
fof(f636, plain, ((e13 = op1(e13, e14)) | ~ (e14 = op1(e13, e13))), inference(cnf_transformation, [], [f12])).
fof(f2466, plain, (~ spl80_3 | spl80_15), inference(avatar_split_clause, [], [f639, f1332, f1282])).
fof(f1282, plain, (spl80_3 <=> (e12 = op1(e14, e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_3])])).
fof(f639, plain, ((e14 = op1(e14, e12)) | ~ (e12 = op1(e14, e14))), inference(cnf_transformation, [], [f12])).
fof(f2450, plain, (spl80_228 | spl80_223 | spl80_218 | spl80_213 | spl80_208), inference(avatar_split_clause, [], [f336, f2193, f2214, f2235, f2256, f2277])).
fof(f336, plain, ((e22 = op2(e21, e24)) | (e22 = op2(e21, e23)) | (e22 = op2(e21, e22)) | (e22 = op2(e21, e21)) | (e22 = op2(e21, e20))), inference(cnf_transformation, [], [f6])).
fof(f6, plain, (((e24 = op2(e24, e24)) | (e24 = op2(e23, e24)) | (e24 = op2(e22, e24)) | (e24 = op2(e21, e24)) | (e24 = op2(e20, e24))) & ((e24 = op2(e24, e24)) | (e24 = op2(e24, e23)) | (e24 = op2(e24, e22)) | (e24 = op2(e24, e21)) | (e24 = op2(e24, e20))) & ((e23 = op2(e24, e24)) | (e23 = op2(e23, e24)) | (e23 = op2(e22, e24)) | (e23 = op2(e21, e24)) | (e23 = op2(e20, e24))) & ((e23 = op2(e24, e24)) | (e23 = op2(e24, e23)) | (e23 = op2(e24, e22)) | (e23 = op2(e24, e21)) | (e23 = op2(e24, e20))) & ((e22 = op2(e24, e24)) | (e22 = op2(e23, e24)) | (e22 = op2(e22, e24)) | (e22 = op2(e21, e24)) | (e22 = op2(e20, e24))) & ((e22 = op2(e24, e24)) | (e22 = op2(e24, e23)) | (e22 = op2(e24, e22)) | (e22 = op2(e24, e21)) | (e22 = op2(e24, e20))) & ((e21 = op2(e24, e24)) | (e21 = op2(e23, e24)) | (e21 = op2(e22, e24)) | (e21 = op2(e21, e24)) | (e21 = op2(e20, e24))) & ((e21 = op2(e24, e24)) | (e21 = op2(e24, e23)) | (e21 = op2(e24, e22)) | (e21 = op2(e24, e21)) | (e21 = op2(e24, e20))) & ((e20 = op2(e24, e24)) | (e20 = op2(e23, e24)) | (e20 = op2(e22, e24)) | (e20 = op2(e21, e24)) | (e20 = op2(e20, e24))) & ((e20 = op2(e24, e24)) | (e20 = op2(e24, e23)) | (e20 = op2(e24, e22)) | (e20 = op2(e24, e21)) | (e20 = op2(e24, e20))) & ((e24 = op2(e24, e23)) | (e24 = op2(e23, e23)) | (e24 = op2(e22, e23)) | (e24 = op2(e21, e23)) | (e24 = op2(e20, e23))) & ((e24 = op2(e23, e24)) | (e24 = op2(e23, e23)) | (e24 = op2(e23, e22)) | (e24 = op2(e23, e21)) | (e24 = op2(e23, e20))) & ((e23 = op2(e24, e23)) | (e23 = op2(e23, e23)) | (e23 = op2(e22, e23)) | (e23 = op2(e21, e23)) | (e23 = op2(e20, e23))) & ((e23 = op2(e23, e24)) | (e23 = op2(e23, e23)) | (e23 = op2(e23, e22)) | (e23 = op2(e23, e21)) | (e23 = op2(e23, e20))) & ((e22 = op2(e24, e23)) | (e22 = op2(e23, e23)) | (e22 = op2(e22, e23)) | (e22 = op2(e21, e23)) | (e22 = op2(e20, e23))) & ((e22 = op2(e23, e24)) | (e22 = op2(e23, e23)) | (e22 = op2(e23, e22)) | (e22 = op2(e23, e21)) | (e22 = op2(e23, e20))) & ((e21 = op2(e24, e23)) | (e21 = op2(e23, e23)) | (e21 = op2(e22, e23)) | (e21 = op2(e21, e23)) | (e21 = op2(e20, e23))) & ((e21 = op2(e23, e24)) | (e21 = op2(e23, e23)) | (e21 = op2(e23, e22)) | (e21 = op2(e23, e21)) | (e21 = op2(e23, e20))) & ((e20 = op2(e24, e23)) | (e20 = op2(e23, e23)) | (e20 = op2(e22, e23)) | (e20 = op2(e21, e23)) | (e20 = op2(e20, e23))) & ((e20 = op2(e23, e24)) | (e20 = op2(e23, e23)) | (e20 = op2(e23, e22)) | (e20 = op2(e23, e21)) | (e20 = op2(e23, e20))) & ((e24 = op2(e24, e22)) | (e24 = op2(e23, e22)) | (e24 = op2(e22, e22)) | (e24 = op2(e21, e22)) | (e24 = op2(e20, e22))) & ((e24 = op2(e22, e24)) | (e24 = op2(e22, e23)) | (e24 = op2(e22, e22)) | (e24 = op2(e22, e21)) | (e24 = op2(e22, e20))) & ((e23 = op2(e24, e22)) | (e23 = op2(e23, e22)) | (e23 = op2(e22, e22)) | (e23 = op2(e21, e22)) | (e23 = op2(e20, e22))) & ((e23 = op2(e22, e24)) | (e23 = op2(e22, e23)) | (e23 = op2(e22, e22)) | (e23 = op2(e22, e21)) | (e23 = op2(e22, e20))) & ((e22 = op2(e24, e22)) | (e22 = op2(e23, e22)) | (e22 = op2(e22, e22)) | (e22 = op2(e21, e22)) | (e22 = op2(e20, e22))) & ((e22 = op2(e22, e24)) | (e22 = op2(e22, e23)) | (e22 = op2(e22, e22)) | (e22 = op2(e22, e21)) | (e22 = op2(e22, e20))) & ((e21 = op2(e24, e22)) | (e21 = op2(e23, e22)) | (e21 = op2(e22, e22)) | (e21 = op2(e21, e22)) | (e21 = op2(e20, e22))) & ((e21 = op2(e22, e24)) | (e21 = op2(e22, e23)) | (e21 = op2(e22, e22)) | (e21 = op2(e22, e21)) | (e21 = op2(e22, e20))) & ((e20 = op2(e24, e22)) | (e20 = op2(e23, e22)) | (e20 = op2(e22, e22)) | (e20 = op2(e21, e22)) | (e20 = op2(e20, e22))) & ((e20 = op2(e22, e24)) | (e20 = op2(e22, e23)) | (e20 = op2(e22, e22)) | (e20 = op2(e22, e21)) | (e20 = op2(e22, e20))) & ((e24 = op2(e24, e21)) | (e24 = op2(e23, e21)) | (e24 = op2(e22, e21)) | (e24 = op2(e21, e21)) | (e24 = op2(e20, e21))) & ((e24 = op2(e21, e24)) | (e24 = op2(e21, e23)) | (e24 = op2(e21, e22)) | (e24 = op2(e21, e21)) | (e24 = op2(e21, e20))) & ((e23 = op2(e24, e21)) | (e23 = op2(e23, e21)) | (e23 = op2(e22, e21)) | (e23 = op2(e21, e21)) | (e23 = op2(e20, e21))) & ((e23 = op2(e21, e24)) | (e23 = op2(e21, e23)) | (e23 = op2(e21, e22)) | (e23 = op2(e21, e21)) | (e23 = op2(e21, e20))) & ((e22 = op2(e24, e21)) | (e22 = op2(e23, e21)) | (e22 = op2(e22, e21)) | (e22 = op2(e21, e21)) | (e22 = op2(e20, e21))) & ((e22 = op2(e21, e24)) | (e22 = op2(e21, e23)) | (e22 = op2(e21, e22)) | (e22 = op2(e21, e21)) | (e22 = op2(e21, e20))) & ((e21 = op2(e24, e21)) | (e21 = op2(e23, e21)) | (e21 = op2(e22, e21)) | (e21 = op2(e21, e21)) | (e21 = op2(e20, e21))) & ((e21 = op2(e21, e24)) | (e21 = op2(e21, e23)) | (e21 = op2(e21, e22)) | (e21 = op2(e21, e21)) | (e21 = op2(e21, e20))) & ((e20 = op2(e24, e21)) | (e20 = op2(e23, e21)) | (e20 = op2(e22, e21)) | (e20 = op2(e21, e21)) | (e20 = op2(e20, e21))) & ((e20 = op2(e21, e24)) | (e20 = op2(e21, e23)) | (e20 = op2(e21, e22)) | (e20 = op2(e21, e21)) | (e20 = op2(e21, e20))) & ((e24 = op2(e24, e20)) | (e24 = op2(e23, e20)) | (e24 = op2(e22, e20)) | (e24 = op2(e21, e20)) | (op2(e20, e20) = e24)) & ((e24 = op2(e20, e24)) | (e24 = op2(e20, e23)) | (e24 = op2(e20, e22)) | (e24 = op2(e20, e21)) | (op2(e20, e20) = e24)) & ((e23 = op2(e24, e20)) | (e23 = op2(e23, e20)) | (e23 = op2(e22, e20)) | (e23 = op2(e21, e20)) | (op2(e20, e20) = e23)) & ((e23 = op2(e20, e24)) | (e23 = op2(e20, e23)) | (e23 = op2(e20, e22)) | (e23 = op2(e20, e21)) | (op2(e20, e20) = e23)) & ((e22 = op2(e24, e20)) | (e22 = op2(e23, e20)) | (e22 = op2(e22, e20)) | (e22 = op2(e21, e20)) | (op2(e20, e20) = e22)) & ((e22 = op2(e20, e24)) | (e22 = op2(e20, e23)) | (e22 = op2(e20, e22)) | (e22 = op2(e20, e21)) | (op2(e20, e20) = e22)) & ((e21 = op2(e24, e20)) | (e21 = op2(e23, e20)) | (e21 = op2(e22, e20)) | (e21 = op2(e21, e20)) | (op2(e20, e20) = e21)) & ((e21 = op2(e20, e24)) | (e21 = op2(e20, e23)) | (e21 = op2(e20, e22)) | (e21 = op2(e20, e21)) | (op2(e20, e20) = e21)) & ((e20 = op2(e24, e20)) | (e20 = op2(e23, e20)) | (e20 = op2(e22, e20)) | (e20 = op2(e21, e20)) | (e20 = op2(e20, e20))) & ((e20 = op2(e20, e24)) | (e20 = op2(e20, e23)) | (e20 = op2(e20, e22)) | (e20 = op2(e20, e21)) | (e20 = op2(e20, e20)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG049+1.p', ax6)).
fof(f2449, plain, (spl80_248 | spl80_223 | spl80_198 | spl80_173 | spl80_148), inference(avatar_split_clause, [], [f337, f1941, f2046, f2151, f2256, f2361])).
fof(f337, plain, ((e22 = op2(e24, e21)) | (e22 = op2(e23, e21)) | (e22 = op2(e22, e21)) | (e22 = op2(e21, e21)) | (e22 = op2(e20, e21))), inference(cnf_transformation, [], [f6])).
fof(f2448, plain, (spl80_229 | spl80_224 | spl80_219 | spl80_214 | spl80_209), inference(avatar_split_clause, [], [f338, f2197, f2218, f2239, f2260, f2281])).
fof(f338, plain, ((e23 = op2(e21, e24)) | (e23 = op2(e21, e23)) | (e23 = op2(e21, e22)) | (e23 = op2(e21, e21)) | (e23 = op2(e21, e20))), inference(cnf_transformation, [], [f6])).
fof(f2447, plain, (spl80_249 | spl80_224 | spl80_199 | spl80_174 | spl80_149), inference(avatar_split_clause, [], [f339, f1945, f2050, f2155, f2260, f2365])).
fof(f339, plain, ((e23 = op2(e24, e21)) | (e23 = op2(e23, e21)) | (e23 = op2(e22, e21)) | (e23 = op2(e21, e21)) | (e23 = op2(e20, e21))), inference(cnf_transformation, [], [f6])).
fof(f2446, plain, (spl80_230 | spl80_225 | spl80_220 | spl80_215 | spl80_210), inference(avatar_split_clause, [], [f340, f2201, f2222, f2243, f2264, f2285])).
fof(f340, plain, ((e24 = op2(e21, e24)) | (e24 = op2(e21, e23)) | (e24 = op2(e21, e22)) | (e24 = op2(e21, e21)) | (e24 = op2(e21, e20))), inference(cnf_transformation, [], [f6])).
fof(f2444, plain, (spl80_201 | spl80_196 | spl80_191 | spl80_186 | spl80_181), inference(avatar_split_clause, [], [f342, f2080, f2101, f2122, f2143, f2164])).
fof(f342, plain, ((e20 = op2(e22, e24)) | (e20 = op2(e22, e23)) | (e20 = op2(e22, e22)) | (e20 = op2(e22, e21)) | (e20 = op2(e22, e20))), inference(cnf_transformation, [], [f6])).
fof(f2436, plain, (spl80_205 | spl80_200 | spl80_195 | spl80_190 | spl80_185), inference(avatar_split_clause, [], [f350, f2096, f2117, f2138, f2159, f2180])).
fof(f350, plain, ((e24 = op2(e22, e24)) | (e24 = op2(e22, e23)) | (e24 = op2(e22, e22)) | (e24 = op2(e22, e21)) | (e24 = op2(e22, e20))), inference(cnf_transformation, [], [f6])).
fof(f2435, plain, (spl80_245 | spl80_220 | spl80_195 | spl80_170 | spl80_145), inference(avatar_split_clause, [], [f351, f1928, f2033, f2138, f2243, f2348])).
fof(f351, plain, ((e24 = op2(e24, e22)) | (e24 = op2(e23, e22)) | (e24 = op2(e22, e22)) | (e24 = op2(e21, e22)) | (e24 = op2(e20, e22))), inference(cnf_transformation, [], [f6])).
fof(f2433, plain, (spl80_236 | spl80_211 | spl80_186 | spl80_161 | spl80_136), inference(avatar_split_clause, [], [f353, f1891, f1996, f2101, f2206, f2311])).
fof(f353, plain, ((e20 = op2(e24, e23)) | (e20 = op2(e23, e23)) | (e20 = op2(e22, e23)) | (e20 = op2(e21, e23)) | (e20 = op2(e20, e23))), inference(cnf_transformation, [], [f6])).
fof(f2431, plain, (spl80_237 | spl80_212 | spl80_187 | spl80_162 | spl80_137), inference(avatar_split_clause, [], [f355, f1895, f2000, f2105, f2210, f2315])).
fof(f355, plain, ((e21 = op2(e24, e23)) | (e21 = op2(e23, e23)) | (e21 = op2(e22, e23)) | (e21 = op2(e21, e23)) | (e21 = op2(e20, e23))), inference(cnf_transformation, [], [f6])).
fof(f2429, plain, (spl80_238 | spl80_213 | spl80_188 | spl80_163 | spl80_138), inference(avatar_split_clause, [], [f357, f1899, f2004, f2109, f2214, f2319])).
fof(f357, plain, ((e22 = op2(e24, e23)) | (e22 = op2(e23, e23)) | (e22 = op2(e22, e23)) | (e22 = op2(e21, e23)) | (e22 = op2(e20, e23))), inference(cnf_transformation, [], [f6])).
fof(f2423, plain, (spl80_231 | spl80_206 | spl80_181 | spl80_156 | spl80_131), inference(avatar_split_clause, [], [f363, f1870, f1975, f2080, f2185, f2290])).
fof(f363, plain, ((e20 = op2(e24, e24)) | (e20 = op2(e23, e24)) | (e20 = op2(e22, e24)) | (e20 = op2(e21, e24)) | (e20 = op2(e20, e24))), inference(cnf_transformation, [], [f6])).
fof(f2421, plain, (spl80_232 | spl80_207 | spl80_182 | spl80_157 | spl80_132), inference(avatar_split_clause, [], [f365, f1874, f1979, f2084, f2189, f2294])).
fof(f365, plain, ((e21 = op2(e24, e24)) | (e21 = op2(e23, e24)) | (e21 = op2(e22, e24)) | (e21 = op2(e21, e24)) | (e21 = op2(e20, e24))), inference(cnf_transformation, [], [f6])).
fof(f2414, plain, (spl80_256 | spl80_257 | spl80_258 | spl80_259 | spl80_260), inference(avatar_split_clause, [], [f321, f2411, f2407, f2403, f2399, f2395])).
fof(f321, plain, ((e24 = unit2) | (e23 = unit2) | (e22 = unit2) | (e21 = unit2) | (e20 = unit2)), inference(cnf_transformation, [], [f5])).
fof(f1852, plain, (spl80_99 | spl80_94 | spl80_89 | spl80_84 | spl80_79), inference(avatar_split_clause, [], [f252, f1601, f1622, f1643, f1664, f1685])).
fof(f252, plain, ((e13 = op1(e11, e14)) | (e13 = op1(e11, e13)) | (e13 = op1(e11, e12)) | (e13 = op1(e11, e11)) | (e13 = op1(e11, e10))), inference(cnf_transformation, [], [f3])).
fof(f3, plain, (((e14 = op1(e14, e14)) | (e14 = op1(e13, e14)) | (e14 = op1(e12, e14)) | (e14 = op1(e11, e14)) | (e14 = op1(e10, e14))) & ((e14 = op1(e14, e14)) | (e14 = op1(e14, e13)) | (e14 = op1(e14, e12)) | (e14 = op1(e14, e11)) | (e14 = op1(e14, e10))) & ((e13 = op1(e14, e14)) | (e13 = op1(e13, e14)) | (e13 = op1(e12, e14)) | (e13 = op1(e11, e14)) | (e13 = op1(e10, e14))) & ((e13 = op1(e14, e14)) | (e13 = op1(e14, e13)) | (e13 = op1(e14, e12)) | (e13 = op1(e14, e11)) | (e13 = op1(e14, e10))) & ((e12 = op1(e14, e14)) | (e12 = op1(e13, e14)) | (e12 = op1(e12, e14)) | (e12 = op1(e11, e14)) | (e12 = op1(e10, e14))) & ((e12 = op1(e14, e14)) | (e12 = op1(e14, e13)) | (e12 = op1(e14, e12)) | (e12 = op1(e14, e11)) | (e12 = op1(e14, e10))) & ((e11 = op1(e14, e14)) | (e11 = op1(e13, e14)) | (e11 = op1(e12, e14)) | (e11 = op1(e11, e14)) | (e11 = op1(e10, e14))) & ((e11 = op1(e14, e14)) | (e11 = op1(e14, e13)) | (e11 = op1(e14, e12)) | (e11 = op1(e14, e11)) | (e11 = op1(e14, e10))) & ((e10 = op1(e14, e14)) | (e10 = op1(e13, e14)) | (e10 = op1(e12, e14)) | (e10 = op1(e11, e14)) | (e10 = op1(e10, e14))) & ((e10 = op1(e14, e14)) | (e10 = op1(e14, e13)) | (e10 = op1(e14, e12)) | (e10 = op1(e14, e11)) | (e10 = op1(e14, e10))) & ((e14 = op1(e14, e13)) | (e14 = op1(e13, e13)) | (e14 = op1(e12, e13)) | (e14 = op1(e11, e13)) | (e14 = op1(e10, e13))) & ((e14 = op1(e13, e14)) | (e14 = op1(e13, e13)) | (e14 = op1(e13, e12)) | (e14 = op1(e13, e11)) | (e14 = op1(e13, e10))) & ((e13 = op1(e14, e13)) | (e13 = op1(e13, e13)) | (e13 = op1(e12, e13)) | (e13 = op1(e11, e13)) | (e13 = op1(e10, e13))) & ((e13 = op1(e13, e14)) | (e13 = op1(e13, e13)) | (e13 = op1(e13, e12)) | (e13 = op1(e13, e11)) | (e13 = op1(e13, e10))) & ((e12 = op1(e14, e13)) | (e12 = op1(e13, e13)) | (e12 = op1(e12, e13)) | (e12 = op1(e11, e13)) | (e12 = op1(e10, e13))) & ((e12 = op1(e13, e14)) | (e12 = op1(e13, e13)) | (e12 = op1(e13, e12)) | (e12 = op1(e13, e11)) | (e12 = op1(e13, e10))) & ((e11 = op1(e14, e13)) | (e11 = op1(e13, e13)) | (e11 = op1(e12, e13)) | (e11 = op1(e11, e13)) | (e11 = op1(e10, e13))) & ((e11 = op1(e13, e14)) | (e11 = op1(e13, e13)) | (e11 = op1(e13, e12)) | (e11 = op1(e13, e11)) | (e11 = op1(e13, e10))) & ((e10 = op1(e14, e13)) | (e10 = op1(e13, e13)) | (e10 = op1(e12, e13)) | (e10 = op1(e11, e13)) | (e10 = op1(e10, e13))) & ((e10 = op1(e13, e14)) | (e10 = op1(e13, e13)) | (e10 = op1(e13, e12)) | (e10 = op1(e13, e11)) | (e10 = op1(e13, e10))) & ((e14 = op1(e14, e12)) | (e14 = op1(e13, e12)) | (e14 = op1(e12, e12)) | (e14 = op1(e11, e12)) | (e14 = op1(e10, e12))) & ((e14 = op1(e12, e14)) | (e14 = op1(e12, e13)) | (e14 = op1(e12, e12)) | (e14 = op1(e12, e11)) | (e14 = op1(e12, e10))) & ((e13 = op1(e14, e12)) | (e13 = op1(e13, e12)) | (e13 = op1(e12, e12)) | (e13 = op1(e11, e12)) | (e13 = op1(e10, e12))) & ((e13 = op1(e12, e14)) | (e13 = op1(e12, e13)) | (e13 = op1(e12, e12)) | (e13 = op1(e12, e11)) | (e13 = op1(e12, e10))) & ((e12 = op1(e14, e12)) | (e12 = op1(e13, e12)) | (e12 = op1(e12, e12)) | (e12 = op1(e11, e12)) | (e12 = op1(e10, e12))) & ((e12 = op1(e12, e14)) | (e12 = op1(e12, e13)) | (e12 = op1(e12, e12)) | (e12 = op1(e12, e11)) | (e12 = op1(e12, e10))) & ((e11 = op1(e14, e12)) | (e11 = op1(e13, e12)) | (e11 = op1(e12, e12)) | (e11 = op1(e11, e12)) | (e11 = op1(e10, e12))) & ((e11 = op1(e12, e14)) | (e11 = op1(e12, e13)) | (e11 = op1(e12, e12)) | (e11 = op1(e12, e11)) | (e11 = op1(e12, e10))) & ((e10 = op1(e14, e12)) | (e10 = op1(e13, e12)) | (e10 = op1(e12, e12)) | (e10 = op1(e11, e12)) | (e10 = op1(e10, e12))) & ((e10 = op1(e12, e14)) | (e10 = op1(e12, e13)) | (e10 = op1(e12, e12)) | (e10 = op1(e12, e11)) | (e10 = op1(e12, e10))) & ((e14 = op1(e14, e11)) | (e14 = op1(e13, e11)) | (e14 = op1(e12, e11)) | (e14 = op1(e11, e11)) | (e14 = op1(e10, e11))) & ((e14 = op1(e11, e14)) | (e14 = op1(e11, e13)) | (e14 = op1(e11, e12)) | (e14 = op1(e11, e11)) | (e14 = op1(e11, e10))) & ((e13 = op1(e14, e11)) | (e13 = op1(e13, e11)) | (e13 = op1(e12, e11)) | (e13 = op1(e11, e11)) | (e13 = op1(e10, e11))) & ((e13 = op1(e11, e14)) | (e13 = op1(e11, e13)) | (e13 = op1(e11, e12)) | (e13 = op1(e11, e11)) | (e13 = op1(e11, e10))) & ((e12 = op1(e14, e11)) | (e12 = op1(e13, e11)) | (e12 = op1(e12, e11)) | (e12 = op1(e11, e11)) | (e12 = op1(e10, e11))) & ((e12 = op1(e11, e14)) | (e12 = op1(e11, e13)) | (e12 = op1(e11, e12)) | (e12 = op1(e11, e11)) | (e12 = op1(e11, e10))) & ((e11 = op1(e14, e11)) | (e11 = op1(e13, e11)) | (e11 = op1(e12, e11)) | (e11 = op1(e11, e11)) | (e11 = op1(e10, e11))) & ((e11 = op1(e11, e14)) | (e11 = op1(e11, e13)) | (e11 = op1(e11, e12)) | (e11 = op1(e11, e11)) | (e11 = op1(e11, e10))) & ((e10 = op1(e14, e11)) | (e10 = op1(e13, e11)) | (e10 = op1(e12, e11)) | (e10 = op1(e11, e11)) | (e10 = op1(e10, e11))) & ((e10 = op1(e11, e14)) | (e10 = op1(e11, e13)) | (e10 = op1(e11, e12)) | (e10 = op1(e11, e11)) | (e10 = op1(e11, e10))) & ((e14 = op1(e14, e10)) | (e14 = op1(e13, e10)) | (e14 = op1(e12, e10)) | (e14 = op1(e11, e10)) | (op1(e10, e10) = e14)) & ((e14 = op1(e10, e14)) | (e14 = op1(e10, e13)) | (e14 = op1(e10, e12)) | (e14 = op1(e10, e11)) | (op1(e10, e10) = e14)) & ((e13 = op1(e14, e10)) | (e13 = op1(e13, e10)) | (e13 = op1(e12, e10)) | (e13 = op1(e11, e10)) | (op1(e10, e10) = e13)) & ((e13 = op1(e10, e14)) | (e13 = op1(e10, e13)) | (e13 = op1(e10, e12)) | (e13 = op1(e10, e11)) | (op1(e10, e10) = e13)) & ((e12 = op1(e14, e10)) | (e12 = op1(e13, e10)) | (e12 = op1(e12, e10)) | (e12 = op1(e11, e10)) | (op1(e10, e10) = e12)) & ((e12 = op1(e10, e14)) | (e12 = op1(e10, e13)) | (e12 = op1(e10, e12)) | (e12 = op1(e10, e11)) | (op1(e10, e10) = e12)) & ((e11 = op1(e14, e10)) | (e11 = op1(e13, e10)) | (e11 = op1(e12, e10)) | (e11 = op1(e11, e10)) | (op1(e10, e10) = e11)) & ((e11 = op1(e10, e14)) | (e11 = op1(e10, e13)) | (e11 = op1(e10, e12)) | (e11 = op1(e10, e11)) | (op1(e10, e10) = e11)) & ((e10 = op1(e14, e10)) | (e10 = op1(e13, e10)) | (e10 = op1(e12, e10)) | (e10 = op1(e11, e10)) | (e10 = op1(e10, e10))) & ((e10 = op1(e10, e14)) | (e10 = op1(e10, e13)) | (e10 = op1(e10, e12)) | (e10 = op1(e10, e11)) | (e10 = op1(e10, e10)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG049+1.p', ax3)).
fof(f1849, plain, (spl80_120 | spl80_95 | spl80_70 | spl80_45 | spl80_20), inference(avatar_split_clause, [], [f255, f1353, f1458, f1563, f1668, f1773])).
fof(f255, plain, ((e14 = op1(e14, e11)) | (e14 = op1(e13, e11)) | (e14 = op1(e12, e11)) | (e14 = op1(e11, e11)) | (e14 = op1(e10, e11))), inference(cnf_transformation, [], [f3])).
fof(f1848, plain, (spl80_71 | spl80_66 | spl80_61 | spl80_56 | spl80_51), inference(avatar_split_clause, [], [f256, f1484, f1505, f1526, f1547, f1568])).
fof(f256, plain, ((e10 = op1(e12, e14)) | (e10 = op1(e12, e13)) | (e10 = op1(e12, e12)) | (e10 = op1(e12, e11)) | (e10 = op1(e12, e10))), inference(cnf_transformation, [], [f3])).
fof(f1846, plain, (spl80_72 | spl80_67 | spl80_62 | spl80_57 | spl80_52), inference(avatar_split_clause, [], [f258, f1488, f1509, f1530, f1551, f1572])).
fof(f258, plain, ((e11 = op1(e12, e14)) | (e11 = op1(e12, e13)) | (e11 = op1(e12, e12)) | (e11 = op1(e12, e11)) | (e11 = op1(e12, e10))), inference(cnf_transformation, [], [f3])).
fof(f1839, plain, (spl80_115 | spl80_90 | spl80_65 | spl80_40 | spl80_15), inference(avatar_split_clause, [], [f265, f1332, f1437, f1542, f1647, f1752])).
fof(f265, plain, ((e14 = op1(e14, e12)) | (e14 = op1(e13, e12)) | (e14 = op1(e12, e12)) | (e14 = op1(e11, e12)) | (e14 = op1(e10, e12))), inference(cnf_transformation, [], [f3])).
fof(f1838, plain, (spl80_46 | spl80_41 | spl80_36 | spl80_31 | spl80_26), inference(avatar_split_clause, [], [f266, f1379, f1400, f1421, f1442, f1463])).
fof(f266, plain, ((e10 = op1(e13, e14)) | (e10 = op1(e13, e13)) | (e10 = op1(e13, e12)) | (e10 = op1(e13, e11)) | (e10 = op1(e13, e10))), inference(cnf_transformation, [], [f3])).
fof(f1834, plain, (spl80_48 | spl80_43 | spl80_38 | spl80_33 | spl80_28), inference(avatar_split_clause, [], [f270, f1387, f1408, f1429, f1450, f1471])).
fof(f270, plain, ((e12 = op1(e13, e14)) | (e12 = op1(e13, e13)) | (e12 = op1(e13, e12)) | (e12 = op1(e13, e11)) | (e12 = op1(e13, e10))), inference(cnf_transformation, [], [f3])).
fof(f1833, plain, (spl80_108 | spl80_83 | spl80_58 | spl80_33 | spl80_8), inference(avatar_split_clause, [], [f271, f1303, f1408, f1513, f1618, f1723])).
fof(f271, plain, ((e12 = op1(e14, e13)) | (e12 = op1(e13, e13)) | (e12 = op1(e12, e13)) | (e12 = op1(e11, e13)) | (e12 = op1(e10, e13))), inference(cnf_transformation, [], [f3])).
fof(f1829, plain, (spl80_110 | spl80_85 | spl80_60 | spl80_35 | spl80_10), inference(avatar_split_clause, [], [f275, f1311, f1416, f1521, f1626, f1731])).
fof(f275, plain, ((e14 = op1(e14, e13)) | (e14 = op1(e13, e13)) | (e14 = op1(e12, e13)) | (e14 = op1(e11, e13)) | (e14 = op1(e10, e13))), inference(cnf_transformation, [], [f3])).
fof(f1827, plain, (spl80_101 | spl80_76 | spl80_51 | spl80_26 | spl80_1), inference(avatar_split_clause, [], [f277, f1274, f1379, f1484, f1589, f1694])).
fof(f277, plain, ((e10 = op1(e14, e14)) | (e10 = op1(e13, e14)) | (e10 = op1(e12, e14)) | (e10 = op1(e11, e14)) | (e10 = op1(e10, e14))), inference(cnf_transformation, [], [f3])).
fof(f1825, plain, (spl80_102 | spl80_77 | spl80_52 | spl80_27 | spl80_2), inference(avatar_split_clause, [], [f279, f1278, f1383, f1488, f1593, f1698])).
fof(f279, plain, ((e11 = op1(e14, e14)) | (e11 = op1(e13, e14)) | (e11 = op1(e12, e14)) | (e11 = op1(e11, e14)) | (e11 = op1(e10, e14))), inference(cnf_transformation, [], [f3])).
fof(f1823, plain, (spl80_103 | spl80_78 | spl80_53 | spl80_28 | spl80_3), inference(avatar_split_clause, [], [f281, f1282, f1387, f1492, f1597, f1702])).
fof(f281, plain, ((e12 = op1(e14, e14)) | (e12 = op1(e13, e14)) | (e12 = op1(e12, e14)) | (e12 = op1(e11, e14)) | (e12 = op1(e10, e14))), inference(cnf_transformation, [], [f3])).
fof(f1818, plain, (spl80_126 | spl80_127 | spl80_128 | spl80_129 | spl80_130), inference(avatar_split_clause, [], [f235, f1815, f1811, f1807, f1803, f1799])).
fof(f235, plain, ((e14 = unit1) | (e13 = unit1) | (e12 = unit1) | (e11 = unit1) | (e10 = unit1)), inference(cnf_transformation, [], [f2])).
fof(f1356, plain, (spl80_16 | spl80_17 | spl80_18 | spl80_19 | spl80_20), inference(avatar_split_clause, [], [f221, f1353, f1349, f1345, f1341, f1337])).
fof(f221, plain, ((e14 = op1(e14, e11)) | (e13 = op1(e14, e11)) | (e12 = op1(e14, e11)) | (e11 = op1(e14, e11)) | (e10 = op1(e14, e11))), inference(cnf_transformation, [], [f1])).
fof(f1, plain, (((e14 = op1(e14, e14)) | (e13 = op1(e14, e14)) | (e12 = op1(e14, e14)) | (e11 = op1(e14, e14)) | (e10 = op1(e14, e14))) & ((e14 = op1(e14, e13)) | (e13 = op1(e14, e13)) | (e12 = op1(e14, e13)) | (e11 = op1(e14, e13)) | (e10 = op1(e14, e13))) & ((e14 = op1(e14, e12)) | (e13 = op1(e14, e12)) | (e12 = op1(e14, e12)) | (e11 = op1(e14, e12)) | (e10 = op1(e14, e12))) & ((e14 = op1(e14, e11)) | (e13 = op1(e14, e11)) | (e12 = op1(e14, e11)) | (e11 = op1(e14, e11)) | (e10 = op1(e14, e11))) & ((e14 = op1(e14, e10)) | (e13 = op1(e14, e10)) | (e12 = op1(e14, e10)) | (e11 = op1(e14, e10)) | (e10 = op1(e14, e10))) & ((e14 = op1(e13, e14)) | (e13 = op1(e13, e14)) | (e12 = op1(e13, e14)) | (e11 = op1(e13, e14)) | (e10 = op1(e13, e14))) & ((e14 = op1(e13, e13)) | (e13 = op1(e13, e13)) | (e12 = op1(e13, e13)) | (e11 = op1(e13, e13)) | (e10 = op1(e13, e13))) & ((e14 = op1(e13, e12)) | (e13 = op1(e13, e12)) | (e12 = op1(e13, e12)) | (e11 = op1(e13, e12)) | (e10 = op1(e13, e12))) & ((e14 = op1(e13, e11)) | (e13 = op1(e13, e11)) | (e12 = op1(e13, e11)) | (e11 = op1(e13, e11)) | (e10 = op1(e13, e11))) & ((e14 = op1(e13, e10)) | (e13 = op1(e13, e10)) | (e12 = op1(e13, e10)) | (e11 = op1(e13, e10)) | (e10 = op1(e13, e10))) & ((e14 = op1(e12, e14)) | (e13 = op1(e12, e14)) | (e12 = op1(e12, e14)) | (e11 = op1(e12, e14)) | (e10 = op1(e12, e14))) & ((e14 = op1(e12, e13)) | (e13 = op1(e12, e13)) | (e12 = op1(e12, e13)) | (e11 = op1(e12, e13)) | (e10 = op1(e12, e13))) & ((e14 = op1(e12, e12)) | (e13 = op1(e12, e12)) | (e12 = op1(e12, e12)) | (e11 = op1(e12, e12)) | (e10 = op1(e12, e12))) & ((e14 = op1(e12, e11)) | (e13 = op1(e12, e11)) | (e12 = op1(e12, e11)) | (e11 = op1(e12, e11)) | (e10 = op1(e12, e11))) & ((e14 = op1(e12, e10)) | (e13 = op1(e12, e10)) | (e12 = op1(e12, e10)) | (e11 = op1(e12, e10)) | (e10 = op1(e12, e10))) & ((e14 = op1(e11, e14)) | (e13 = op1(e11, e14)) | (e12 = op1(e11, e14)) | (e11 = op1(e11, e14)) | (e10 = op1(e11, e14))) & ((e14 = op1(e11, e13)) | (e13 = op1(e11, e13)) | (e12 = op1(e11, e13)) | (e11 = op1(e11, e13)) | (e10 = op1(e11, e13))) & ((e14 = op1(e11, e12)) | (e13 = op1(e11, e12)) | (e12 = op1(e11, e12)) | (e11 = op1(e11, e12)) | (e10 = op1(e11, e12))) & ((e14 = op1(e11, e11)) | (e13 = op1(e11, e11)) | (e12 = op1(e11, e11)) | (e11 = op1(e11, e11)) | (e10 = op1(e11, e11))) & ((e14 = op1(e11, e10)) | (e13 = op1(e11, e10)) | (e12 = op1(e11, e10)) | (e11 = op1(e11, e10)) | (e10 = op1(e11, e10))) & ((e14 = op1(e10, e14)) | (e13 = op1(e10, e14)) | (e12 = op1(e10, e14)) | (e11 = op1(e10, e14)) | (e10 = op1(e10, e14))) & ((e14 = op1(e10, e13)) | (e13 = op1(e10, e13)) | (e12 = op1(e10, e13)) | (e11 = op1(e10, e13)) | (e10 = op1(e10, e13))) & ((e14 = op1(e10, e12)) | (e13 = op1(e10, e12)) | (e12 = op1(e10, e12)) | (e11 = op1(e10, e12)) | (e10 = op1(e10, e12))) & ((e14 = op1(e10, e11)) | (e13 = op1(e10, e11)) | (e12 = op1(e10, e11)) | (e11 = op1(e10, e11)) | (e10 = op1(e10, e11))) & ((op1(e10, e10) = e14) | (op1(e10, e10) = e13) | (op1(e10, e10) = e12) | (op1(e10, e10) = e11) | (e10 = op1(e10, e10)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG049+1.p', ax1)).