fof(f16685, plain, $false, inference(avatar_sat_refutation, [], [f664, f685, f706, f727, f748, f769, f790, f811, f832, f853, f874, f895, f916, f937, f958, f979, f1000, f1021, f1042, f1063, f1084, f1105, f1126, f1147, f1168, f1189, f1192, f1193, f1194, f1195, f1196, f1197, f1198, f1199, f1200, f1201, f1204, f1205, f1206, f1207, f1208, f1209, f1210, f1211, f1212, f1213, f1216, f1217, f1218, f1219, f1220, f1221, f1222, f1223, f1224, f1225, f1228, f1229, f1230, f1231, f1232, f1233, f1234, f1235, f1236, f1237, f1239, f1248, f1253, f1262, f1267, f1276, f1281, f1290, f1295, f1304, f1309, f1318, f1323, f1332, f1333, f1342, f1347, f1356, f1361, f1370, f1371, f1384, f1389, f1394, f1395, f1396, f1397, f1398, f1399, f1412, f1417, f1418, f1419, f1424, f1425, f1426, f1427, f1440, f1441, f1446, f1447, f1452, f1453, f1454, f1455, f1468, f1473, f1478, f1479, f1484, f1485, f1486, f1487, f1500, f1505, f1510, f1511, f1516, f1517, f1518, f1519, f1532, f1537, f1550, f1555, f1568, f1601, f1636, f1642, f1643, f1648, f1654, f1655, f1664, f1665, f1674, f1675, f1684, f1703, f1708, f1714, f1715, f1730, f1732, f1741, f1742, f1743, f1752, f1758, f1759, f1764, f1775, f1777, f1778, f1784, f1785, f1791, f1792, f1793, f1794, f1796, f1797, f1798, f1799, f1804, f1805, f1817, f1823, f1825, f1827, f1829, f1830, f1831, f1832, f1835, f1837, f1838, f1844, f1845, f1850, f1851, f1856, f1857, f1862, f1863, f1864, f1867, f1868, f1870, f1871, f1872, f1873, f1874, f1875, f1912, f1914, f1916, f1918, f1935, f1938, f1939, f1959, f1974, f1977, f1980, f1992, f2001, f2019, f2030, f2059, f2065, f2124, f2130, f2175, f2186, f2189, f2192, f2211, f2214, f2217, f2222, f2225, f2239, f2248, f2251, f2253, f2256, f2301, f2308, f2315, f2319, f2332, f2337, f2347, f2349, f2364, f2368, f2387, f2389, f2408, f2419, f2442, f2448, f2458, f2491, f2498, f2546, f2573, f2593, f2598, f2601, f2607, f2612, f2621, f2622, f2636, f2659, f2672, f2678, f2681, f2707, f2718, f2762, f2804, f2811, f2814, f2824, f2833, f2852, f2888, f2951, f2962, f2984, f3000, f3039, f3077, f3085, f3096, f3098, f3103, f3111, f3256, f3286, f3289, f3291, f3311, f3344, f3357, f3373, f3390, f3393, f3427, f3431, f3458, f3530, f3533, f3548, f3565, f3588, f3591, f3603, f3614, f3681, f3693, f3742, f3777, f3796, f3812, f3831, f3834, f3853, f3907, f3943, f3952, f4002, f4018, f4023, f4042, f4063, f4094, f4106, f4149, f4178, f4184, f4187, f4190, f4210, f4226, f4243, f4275, f4285, f4323, f4332, f4334, f4339, f4374, f4375, f4382, f4406, f4442, f4449, f4452, f4459, f4461, f4467, f4478, f4497, f4558, f4565, f4579, f4599, f4607, f4633, f4638, f4649, f4651, f4707, f4727, f4728, f4738, f4749, f4761, f4762, f4779, f4783, f4822, f4825, f4833, f4841, f4845, f4878, f4915, f4928, f4936, f4948, f4965, f4985, f4988, f5032, f5049, f5092, f5150, f5154, f5161, f5176, f5201, f5218, f5219, f5220, f5239, f5250, f5264, f5273, f5285, f5295, f5326, f5327, f5333, f5376, f5406, f5411, f5459, f5474, f5481, f5514, f5531, f5553, f5575, f5597, f5624, f5627, f5648, f5650, f5663, f5683, f5684, f5702, f5752, f5797, f5811, f5814, f5849, f5874, f5876, f5911, f5914, f5922, f5931, f5939, f5996, f6005, f6049, f6093, f6103, f6146, f6148, f6149, f6150, f6159, f6160, f6219, f6222, f6268, f6297, f6306, f6322, f6354, f6468, f6501, f6514, f6516, f6526, f6547, f6556, f6563, f6573, f6597, f6602, f6621, f6640, f6668, f6677, f6684, f6720, f6734, f6766, f6816, f6831, f6851, f6895, f6897, f6905, f6973, f6988, f6996, f7013, f7069, f7096, f7120, f7138, f7150, f7186, f7190, f7199, f7217, f7219, f7263, f7282, f7306, f7309, f7311, f7316, f7335, f7378, f7379, f7398, f7409, f7418, f7421, f7425, f7437, f7446, f7459, f7471, f7502, f7516, f7534, f7563, f7565, f7581, f7610, f7613, f7625, f7636, f7639, f7681, f7700, f7715, f7759, f7770, f7789, f7829, f7837, f7852, f7858, f7870, f7882, f7883, f7891, f7892, f7902, f7903, f7904, f7914, f7943, f7944, f7972, f7974, f7978, f7989, f8015, f8030, f8103, f8104, f8182, f8191, f8199, f8222, f8246, f8277, f8284, f8290, f8389, f8425, f8485, f8511, f8552, f8581, f8594, f8626, f8634, f8638, f8697, f8698, f8711, f8728, f8746, f8758, f8759, f8778, f8809, f8821, f8826, f8845, f8857, f8865, f8900, f8907, f8988, f9025, f9081, f9103, f9139, f9155, f9173, f9175, f9199, f9240, f9241, f9302, f9305, f9369, f9436, f9440, f9462, f9514, f9519, f9523, f9553, f9596, f9605, f9611, f9614, f9627, f9628, f9665, f9666, f9668, f9754, f9771, f9775, f9795, f9819, f9835, f9838, f9867, f9877, f9880, f9882, f9921, f9931, f9959, f9964, f9977, f10004, f10009, f10109, f10122, f10142, f10143, f10144, f10167, f10176, f10245, f10282, f10288, f10313, f10350, f10351, f10355, f10362, f10364, f10374, f10395, f10406, f10411, f10413, f10432, f10443, f10444, f10452, f10453, f10464, f10476, f10520, f10535, f10543, f10547, f10572, f10598, f10617, f10620, f10622, f10624, f10635, f10695, f10698, f10708, f10723, f10731, f10733, f10743, f10785, f10797, f10818, f10879, f10921, f10952, f10967, f10989, f10991, f10993, f10997, f10999, f11001, f11021, f11066, f11074, f11103, f11111, f11125, f11145, f11153, f11178, f11187, f11207, f11214, f11224, f11230, f11239, f11243, f11258, f11260, f11262, f11271, f11278, f11288, f11305, f11306, f11308, f11325, f11346, f11359, f11374, f11382, f11415, f11422, f11424, f11426, f11461, f11511, f11522, f11524, f11544, f11613, f11614, f11649, f11666, f11685, f11781, f11791, f11830, f11839, f11849, f11859, f11865, f11867, f11882, f11903, f11919, f11921, f11922, f11957, f11981, f11986, f11988, f11990, f12007, f12014, f12037, f12038, f12040, f12041, f12042, f12043, f12044, f12045, f12046, f12059, f12072, f12085, f12086, f12105, f12106, f12125, f12126, f12132, f12138, f12157, f12158, f12174, f12187, f12196, f12205, f12208, f12217, f12226, f12255, f12258, f12260, f12263, f12264, f12268, f12270, f12314, f12327, f12329, f12330, f12331, f12340, f12350, f12351, f12365, f12376, f12386, f12394, f12405, f12429, f12434, f12442, f12495, f12500, f12502, f12504, f12507, f12515, f12525, f12530, f12543, f12549, f12563, f12564, f12565, f12566, f12567, f12568, f12570, f12571, f12572, f12593, f12594, f12622, f12623, f12651, f12677, f12688, f12705, f12710, f12720, f12745, f12748, f12777, f12783, f12785, f12791, f12808, f12819, f12855, f12871, f12872, f12881, f12882, f12901, f12902, f12903, f12927, f12937, f12959, f12975, f12977, f12983, f13009, f13033, f13035, f13045, f13046, f13047, f13063, f13075, f13103, f13105, f13126, f13147, f13153, f13155, f13164, f13166, f13167, f13230, f13231, f13233, f13248, f13250, f13262, f13263, f13274, f13276, f13308, f13342, f13354, f13369, f13379, f13384, f13395, f13404, f13412, f13423, f13426, f13447, f13452, f13455, f13497, f13515, f13544, f13546, f13548, f13550, f13558, f13578, f13593, f13596, f13599, f13601, f13630, f13631, f13632, f13641, f13643, f13674, f13675, f13694, f13716, f13723, f13831, f13833, f13837, f13840, f13873, f13898, f13947, f13951, f13953, f13954, f13981, f13982, f13991, f13994, f13998, f14000, f14022, f14023, f14031, f14047, f14056, f14066, f14067, f14068, f14072, f14090, f14101, f14115, f14117, f14127, f14131, f14135, f14137, f14139, f14148, f14150, f14152, f14161, f14163, f14177, f14196, f14197, f14198, f14199, f14200, f14201, f14202, f14203, f14204, f14224, f14225, f14243, f14244, f14254, f14262, f14263, f14270, f14281, f14289, f14300, f14301, f14337, f14346, f14362, f14369, f14376, f14378, f14386, f14392, f14413, f14423, f14429, f14431, f14447, f14449, f14462, f14463, f14464, f14465, f14466, f14474, f14475, f14477, f14498, f14499, f14512, f14524, f14527, f14528, f14545, f14554, f14555, f14566, f14576, f14577, f14585, f14594, f14617, f14627, f14633, f14655, f14659, f14661, f14664, f14684, f14685, f14699, f14719, f14748, f14756, f14765, f14791, f14802, f14840, f14842, f14843, f14854, f14855, f14856, f14889, f14890, f14902, f14918, f14919, f14920, f14921, f14932, f14944, f14945, f14953, f15013, f15014, f15030, f15052, f15059, f15069, f15079, f15084, f15092, f15102, f15124, f15141, f15143, f15147, f15163, f15164, f15165, f15177, f15187, f15193, f15199, f15210, f15211, f15217, f15225, f15229, f15244, f15251, f15263, f15265, f15267, f15269, f15271, f15274, f15277, f15280, f15313, f15348, f15349, f15351, f15353, f15362, f15363, f15374, f15375, f15386, f15405, f15406, f15409, f15419, f15420, f15440, f15441, f15452, f15466, f15475, f15500, f15508, f15522, f15557, f15581, f15652, f15656, f15663, f15681, f15682, f15683, f15687, f15693, f15694, f15701, f15710, f15722, f15729, f15733, f15738, f15752, f15759, f15772, f15773, f15782, f15784, f15786, f15790, f15792, f15793, f15795, f15797, f15820, f15822, f15835, f15836, f15837, f15838, f15840, f15841, f15842, f15843, f15844, f15863, f15864, f15865, f15880, f15881, f15891, f15892, f15905, f15922, f15923, f15938, f15949, f15970, f15971, f15983, f15984, f15995, f15996, f16003, f16008, f16016, f16027, f16034, f16047, f16062, f16070, f16072, f16077, f16098, f16105, f16108, f16113, f16126, f16128, f16130, f16141, f16142, f16143, f16144, f16145, f16146, f16147, f16148, f16150, f16167, f16168, f16169, f16171, f16173, f16183, f16184, f16185, f16186, f16196, f16207, f16222, f16223, f16242, f16250, f16251, f16265, f16267, f16269, f16279, f16297, f16311, f16315, f16332, f16399, f16400, f16422, f16437, f16441, f16455, f16456, f16457, f16459, f16467, f16471, f16478, f16479, f16492, f16493, f16494, f16495, f16498, f16503, f16507, f16521, f16527, f16536, f16540, f16549, f16551, f16572, f16574, f16579, f16601, f16602, f16603, f16605, f16612, f16616, f16623, f16633, f16643, f16649, f16664, f16666, f16670, f16683])).
fof(f16683, plain, (~ spl10_32 | ~ spl10_118 | spl10_202 | ~ spl10_224), inference(avatar_contradiction_clause, [], [f16682])).
fof(f16682, plain, ($false | (~ spl10_32 | ~ spl10_118 | spl10_202 | ~ spl10_224)), inference(subsumption_resolution, [], [f16679, f1138])).
fof(f1138, plain, ((e2 = op(e0, e1)) | ~ spl10_118), inference(avatar_component_clause, [], [f1136])).
fof(f1136, plain, (spl10_118 <=> (e2 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl10_118])])).
fof(f16679, plain, (~ (e2 = op(e0, e1)) | (~ spl10_32 | spl10_202 | ~ spl10_224)), inference(backward_demodulation, [], [f16562, f777])).
fof(f777, plain, ((e1 = op(e3, e3)) | ~ spl10_32), inference(avatar_component_clause, [], [f775])).
fof(f775, plain, (spl10_32 <=> (e1 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl10_32])])).
fof(f16562, plain, (~ (e2 = op(e0, op(e3, e3))) | (spl10_202 | ~ spl10_224)), inference(backward_demodulation, [], [f1595, f1706])).
fof(f1706, plain, ((e0 = op(e3, op(e3, e3))) | ~ spl10_224), inference(avatar_component_clause, [], [f1705])).
fof(f1705, plain, (spl10_224 <=> (e0 = op(e3, op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl10_224])])).
fof(f1595, plain, (~ (e2 = op(op(e3, op(e3, e3)), op(e3, e3))) | spl10_202), inference(avatar_component_clause, [], [f1593])).
fof(f1593, plain, (spl10_202 <=> (e2 = op(op(e3, op(e3, e3)), op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl10_202])])).
fof(f16670, plain, (~ spl10_35 | ~ spl10_40), inference(avatar_split_clause, [], [f16667, f808, f787])).
fof(f787, plain, (spl10_35 <=> (e4 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl10_35])])).
fof(f808, plain, (spl10_40 <=> (e4 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl10_40])])).
fof(f16667, plain, (~ (e4 = op(e3, e3)) | ~ spl10_40), inference(backward_demodulation, [], [f441, f810])).
fof(f810, plain, ((e4 = op(e3, e2)) | ~ spl10_40), inference(avatar_component_clause, [], [f808])).
fof(f441, plain, ~ (op(e3, e2) = op(e3, e3)), inference(cnf_transformation, [], [f4])).
fof(f4, plain, (~ (op(e4, e3) = op(e4, e4)) & ~ (op(e4, e2) = op(e4, e4)) & ~ (op(e4, e2) = op(e4, e3)) & ~ (op(e4, e1) = op(e4, e4)) & ~ (op(e4, e1) = op(e4, e3)) & ~ (op(e4, e1) = op(e4, e2)) & ~ (op(e4, e0) = op(e4, e4)) & ~ (op(e4, e0) = op(e4, e3)) & ~ (op(e4, e0) = op(e4, e2)) & ~ (op(e4, e0) = op(e4, e1)) & ~ (op(e3, e3) = op(e3, e4)) & ~ (op(e3, e2) = op(e3, e4)) & ~ (op(e3, e2) = op(e3, e3)) & ~ (op(e3, e1) = op(e3, e4)) & ~ (op(e3, e1) = op(e3, e3)) & ~ (op(e3, e1) = op(e3, e2)) & ~ (op(e3, e0) = op(e3, e4)) & ~ (op(e3, e0) = op(e3, e3)) & ~ (op(e3, e0) = op(e3, e2)) & ~ (op(e3, e0) = op(e3, e1)) & ~ (op(e2, e3) = op(e2, e4)) & ~ (op(e2, e2) = op(e2, e4)) & ~ (op(e2, e2) = op(e2, e3)) & ~ (op(e2, e1) = op(e2, e4)) & ~ (op(e2, e1) = op(e2, e3)) & ~ (op(e2, e1) = op(e2, e2)) & ~ (op(e2, e0) = op(e2, e4)) & ~ (op(e2, e0) = op(e2, e3)) & ~ (op(e2, e0) = op(e2, e2)) & ~ (op(e2, e0) = op(e2, e1)) & ~ (op(e1, e3) = op(e1, e4)) & ~ (op(e1, e2) = op(e1, e4)) & ~ (op(e1, e2) = op(e1, e3)) & ~ (op(e1, e1) = op(e1, e4)) & ~ (op(e1, e1) = op(e1, e3)) & ~ (op(e1, e1) = op(e1, e2)) & ~ (op(e1, e0) = op(e1, e4)) & ~ (op(e1, e0) = op(e1, e3)) & ~ (op(e1, e0) = op(e1, e2)) & ~ (op(e1, e0) = op(e1, e1)) & ~ (op(e0, e3) = op(e0, e4)) & ~ (op(e0, e2) = op(e0, e4)) & ~ (op(e0, e2) = op(e0, e3)) & ~ (op(e0, e1) = op(e0, e4)) & ~ (op(e0, e1) = op(e0, e3)) & ~ (op(e0, e1) = op(e0, e2)) & ~ (op(e0, e0) = op(e0, e4)) & ~ (op(e0, e0) = op(e0, e3)) & ~ (op(e0, e0) = op(e0, e2)) & ~ (op(e0, e0) = op(e0, e1)) & ~ (op(e3, e4) = op(e4, e4)) & ~ (op(e2, e4) = op(e4, e4)) & ~ (op(e2, e4) = op(e3, e4)) & ~ (op(e1, e4) = op(e4, e4)) & ~ (op(e1, e4) = op(e3, e4)) & ~ (op(e1, e4) = op(e2, e4)) & ~ (op(e0, e4) = op(e4, e4)) & ~ (op(e0, e4) = op(e3, e4)) & ~ (op(e0, e4) = op(e2, e4)) & ~ (op(e0, e4) = op(e1, e4)) & ~ (op(e3, e3) = op(e4, e3)) & ~ (op(e2, e3) = op(e4, e3)) & ~ (op(e2, e3) = op(e3, e3)) & ~ (op(e1, e3) = op(e4, e3)) & ~ (op(e1, e3) = op(e3, e3)) & ~ (op(e1, e3) = op(e2, e3)) & ~ (op(e0, e3) = op(e4, e3)) & ~ (op(e0, e3) = op(e3, e3)) & ~ (op(e0, e3) = op(e2, e3)) & ~ (op(e0, e3) = op(e1, e3)) & ~ (op(e3, e2) = op(e4, e2)) & ~ (op(e2, e2) = op(e4, e2)) & ~ (op(e2, e2) = op(e3, e2)) & ~ (op(e1, e2) = op(e4, e2)) & ~ (op(e1, e2) = op(e3, e2)) & ~ (op(e1, e2) = op(e2, e2)) & ~ (op(e0, e2) = op(e4, e2)) & ~ (op(e0, e2) = op(e3, e2)) & ~ (op(e0, e2) = op(e2, e2)) & ~ (op(e0, e2) = op(e1, e2)) & ~ (op(e3, e1) = op(e4, e1)) & ~ (op(e2, e1) = op(e4, e1)) & ~ (op(e2, e1) = op(e3, e1)) & ~ (op(e1, e1) = op(e4, e1)) & ~ (op(e1, e1) = op(e3, e1)) & ~ (op(e1, e1) = op(e2, e1)) & ~ (op(e0, e1) = op(e4, e1)) & ~ (op(e0, e1) = op(e3, e1)) & ~ (op(e0, e1) = op(e2, e1)) & ~ (op(e0, e1) = op(e1, e1)) & ~ (op(e3, e0) = op(e4, e0)) & ~ (op(e2, e0) = op(e4, e0)) & ~ (op(e2, e0) = op(e3, e0)) & ~ (op(e1, e0) = op(e4, e0)) & ~ (op(e1, e0) = op(e3, e0)) & ~ (op(e1, e0) = op(e2, e0)) & ~ (op(e0, e0) = op(e4, e0)) & ~ (op(e0, e0) = op(e3, e0)) & ~ (op(e0, e0) = op(e2, e0)) & ~ (op(e0, e0) = op(e1, e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax4)).
fof(f16666, plain, (~ spl10_62 | ~ spl10_95 | spl10_243), inference(avatar_contradiction_clause, [], [f16665])).
fof(f16665, plain, ($false | (~ spl10_62 | ~ spl10_95 | spl10_243)), inference(subsumption_resolution, [], [f16662, f1041])).
fof(f1041, plain, ((e4 = op(e1, e1)) | ~ spl10_95), inference(avatar_component_clause, [], [f1039])).
fof(f1039, plain, (spl10_95 <=> (e4 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl10_95])])).
fof(f16662, plain, (~ (e4 = op(e1, e1)) | (~ spl10_62 | spl10_243)), inference(backward_demodulation, [], [f1849, f903])).
fof(f903, plain, ((e1 = op(e2, e2)) | ~ spl10_62), inference(avatar_component_clause, [], [f901])).
fof(f901, plain, (spl10_62 <=> (e1 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl10_62])])).
fof(f1849, plain, (~ (e4 = op(op(e2, e2), op(e2, e2))) | spl10_243), inference(avatar_component_clause, [], [f1847])).
fof(f1847, plain, (spl10_243 <=> (e4 = op(op(e2, e2), op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl10_243])])).
fof(f16664, plain, (~ spl10_37 | ~ spl10_62), inference(avatar_split_clause, [], [f16651, f901, f796])).
fof(f796, plain, (spl10_37 <=> (e1 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl10_37])])).
fof(f16651, plain, (~ (e1 = op(e3, e2)) | ~ spl10_62), inference(backward_demodulation, [], [f381, f903])).
fof(f381, plain, ~ (op(e2, e2) = op(e3, e2)), inference(cnf_transformation, [], [f4])).
fof(f16649, plain, (~ spl10_64 | ~ spl10_69), inference(avatar_split_clause, [], [f16645, f930, f909])).
fof(f909, plain, (spl10_64 <=> (e3 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl10_64])])).
fof(f930, plain, (spl10_69 <=> (e3 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl10_69])])).
fof(f16645, plain, (~ (e3 = op(e2, e2)) | ~ spl10_69), inference(backward_demodulation, [], [f428, f932])).
fof(f932, plain, ((e3 = op(e2, e1)) | ~ spl10_69), inference(avatar_component_clause, [], [f930])).
fof(f428, plain, ~ (op(e2, e1) = op(e2, e2)), inference(cnf_transformation, [], [f4])).
fof(f16643, plain, (~ spl10_70 | ~ spl10_75), inference(avatar_split_clause, [], [f16639, f955, f934])).
fof(f934, plain, (spl10_70 <=> (e4 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl10_70])])).
fof(f955, plain, (spl10_75 <=> (e4 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl10_75])])).
fof(f16639, plain, (~ (e4 = op(e2, e1)) | ~ spl10_75), inference(backward_demodulation, [], [f424, f957])).
fof(f957, plain, ((e4 = op(e2, e0)) | ~ spl10_75), inference(avatar_component_clause, [], [f955])).
fof(f424, plain, ~ (op(e2, e0) = op(e2, e1)), inference(cnf_transformation, [], [f4])).
fof(f16633, plain, (~ spl10_70 | ~ spl10_95), inference(avatar_split_clause, [], [f16624, f1039, f934])).
fof(f16624, plain, (~ (e4 = op(e2, e1)) | ~ spl10_95), inference(backward_demodulation, [], [f368, f1041])).
fof(f368, plain, ~ (op(e1, e1) = op(e2, e1)), inference(cnf_transformation, [], [f4])).
fof(f16623, plain, (~ spl10_94 | ~ spl10_99), inference(avatar_split_clause, [], [f16619, f1056, f1035])).
fof(f1035, plain, (spl10_94 <=> (e3 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl10_94])])).
fof(f1056, plain, (spl10_99 <=> (e3 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl10_99])])).
fof(f16619, plain, (~ (e3 = op(e1, e1)) | ~ spl10_99), inference(backward_demodulation, [], [f414, f1058])).
fof(f1058, plain, ((e3 = op(e1, e0)) | ~ spl10_99), inference(avatar_component_clause, [], [f1056])).
fof(f414, plain, ~ (op(e1, e0) = op(e1, e1)), inference(cnf_transformation, [], [f4])).
fof(f16616, plain, (~ spl10_35 | ~ spl10_110), inference(avatar_split_clause, [], [f16613, f1102, f787])).
fof(f1102, plain, (spl10_110 <=> (e4 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl10_110])])).
fof(f16613, plain, (~ (e4 = op(e3, e3)) | ~ spl10_110), inference(backward_demodulation, [], [f386, f1104])).
fof(f1104, plain, ((e4 = op(e0, e3)) | ~ spl10_110), inference(avatar_component_clause, [], [f1102])).
fof(f386, plain, ~ (op(e0, e3) = op(e3, e3)), inference(cnf_transformation, [], [f4])).
fof(f16612, plain, (~ spl10_64 | ~ spl10_114), inference(avatar_split_clause, [], [f16606, f1119, f909])).
fof(f1119, plain, (spl10_114 <=> (e3 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl10_114])])).
fof(f16606, plain, (~ (e3 = op(e2, e2)) | ~ spl10_114), inference(backward_demodulation, [], [f375, f1121])).
fof(f1121, plain, ((e3 = op(e0, e2)) | ~ spl10_114), inference(avatar_component_clause, [], [f1119])).
fof(f375, plain, ~ (op(e0, e2) = op(e2, e2)), inference(cnf_transformation, [], [f4])).
fof(f16605, plain, (~ spl10_69 | ~ spl10_122 | spl10_226 | ~ spl10_227), inference(avatar_split_clause, [], [f16600, f1721, f1717, f1153, f930])).
fof(f1153, plain, (spl10_122 <=> (op(e0, e0) = e1)), introduced(avatar_definition, [new_symbols(naming, [spl10_122])])).
fof(f1717, plain, (spl10_226 <=> (e3 = op(op(e0, op(e0, e0)), op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl10_226])])).
fof(f1721, plain, (spl10_227 <=> (e2 = op(e0, op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl10_227])])).
fof(f16600, plain, (~ (e3 = op(e2, e1)) | (~ spl10_122 | spl10_226 | ~ spl10_227)), inference(backward_demodulation, [], [f16560, f1155])).
fof(f1155, plain, ((op(e0, e0) = e1) | ~ spl10_122), inference(avatar_component_clause, [], [f1153])).
fof(f16560, plain, (~ (e3 = op(e2, op(e0, e0))) | (spl10_226 | ~ spl10_227)), inference(forward_demodulation, [], [f1719, f1722])).
fof(f1722, plain, ((e2 = op(e0, op(e0, e0))) | ~ spl10_227), inference(avatar_component_clause, [], [f1721])).
fof(f1719, plain, (~ (e3 = op(op(e0, op(e0, e0)), op(e0, e0))) | spl10_226), inference(avatar_component_clause, [], [f1717])).
fof(f16603, plain, (~ spl10_94 | ~ spl10_122 | spl10_241), inference(avatar_split_clause, [], [f16597, f1819, f1153, f1035])).
fof(f1819, plain, (spl10_241 <=> (e3 = op(op(e0, e0), op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl10_241])])).
fof(f16597, plain, (~ (e3 = op(e1, e1)) | (~ spl10_122 | spl10_241)), inference(backward_demodulation, [], [f1821, f1155])).
fof(f1821, plain, (~ (e3 = op(op(e0, e0), op(e0, e0))) | spl10_241), inference(avatar_component_clause, [], [f1819])).
fof(f16602, plain, (~ spl10_107 | ~ spl10_122), inference(avatar_split_clause, [], [f16589, f1153, f1090])).
fof(f1090, plain, (spl10_107 <=> (e1 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl10_107])])).
fof(f16589, plain, (~ (e1 = op(e0, e3)) | ~ spl10_122), inference(backward_demodulation, [], [f406, f1155])).
fof(f406, plain, ~ (op(e0, e0) = op(e0, e3)), inference(cnf_transformation, [], [f4])).
fof(f16601, plain, (~ spl10_72 | ~ spl10_122), inference(avatar_split_clause, [], [f16587, f1153, f943])).
fof(f943, plain, (spl10_72 <=> (e1 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl10_72])])).
fof(f16587, plain, (~ (e1 = op(e2, e0)) | ~ spl10_122), inference(backward_demodulation, [], [f355, f1155])).
fof(f355, plain, ~ (op(e0, e0) = op(e2, e0)), inference(cnf_transformation, [], [f4])).
fof(f16579, plain, (~ spl10_35 | ~ spl10_29 | spl10_163), inference(avatar_split_clause, [], [f16578, f1391, f762, f787])).
fof(f762, plain, (spl10_29 <=> (e3 = op(e3, e4))), introduced(avatar_definition, [new_symbols(naming, [spl10_29])])).
fof(f1391, plain, (spl10_163 <=> (e4 = op(e3, op(e3, e4)))), introduced(avatar_definition, [new_symbols(naming, [spl10_163])])).
fof(f16578, plain, (~ (e4 = op(e3, e3)) | (~ spl10_29 | spl10_163)), inference(forward_demodulation, [], [f1393, f764])).
fof(f764, plain, ((e3 = op(e3, e4)) | ~ spl10_29), inference(avatar_component_clause, [], [f762])).
fof(f1393, plain, (~ (e4 = op(e3, op(e3, e4))) | spl10_163), inference(avatar_component_clause, [], [f1391])).
fof(f16574, plain, (spl10_99 | ~ spl10_41 | ~ spl10_175), inference(avatar_split_clause, [], [f16573, f1461, f813, f1056])).
fof(f813, plain, (spl10_41 <=> (e0 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl10_41])])).
fof(f1461, plain, (spl10_175 <=> (e3 = op(e1, op(e3, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl10_175])])).
fof(f16573, plain, ((e3 = op(e1, e0)) | (~ spl10_41 | ~ spl10_175)), inference(forward_demodulation, [], [f1463, f815])).
fof(f815, plain, ((e0 = op(e3, e1)) | ~ spl10_41), inference(avatar_component_clause, [], [f813])).
fof(f1463, plain, ((e3 = op(e1, op(e3, e1))) | ~ spl10_175), inference(avatar_component_clause, [], [f1461])).
fof(f16572, plain, (spl10_114 | ~ spl10_48 | ~ spl10_181), inference(avatar_split_clause, [], [f16571, f1493, f842, f1119])).
fof(f842, plain, (spl10_48 <=> (e2 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl10_48])])).
fof(f1493, plain, (spl10_181 <=> (e3 = op(e0, op(e3, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl10_181])])).
fof(f16571, plain, ((e3 = op(e0, e2)) | (~ spl10_48 | ~ spl10_181)), inference(forward_demodulation, [], [f1495, f844])).
fof(f844, plain, ((e2 = op(e3, e0)) | ~ spl10_48), inference(avatar_component_clause, [], [f842])).
fof(f1495, plain, ((e3 = op(e0, op(e3, e0))) | ~ spl10_181), inference(avatar_component_clause, [], [f1493])).
fof(f16551, plain, (~ spl10_108 | ~ spl10_83), inference(avatar_split_clause, [], [f16550, f989, f1094])).
fof(f1094, plain, (spl10_108 <=> (e2 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl10_108])])).
fof(f989, plain, (spl10_83 <=> (e2 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl10_83])])).
fof(f16550, plain, (~ (e2 = op(e0, e3)) | ~ spl10_83), inference(forward_demodulation, [], [f384, f991])).
fof(f991, plain, ((e2 = op(e1, e3)) | ~ spl10_83), inference(avatar_component_clause, [], [f989])).
fof(f384, plain, ~ (op(e0, e3) = op(e1, e3)), inference(cnf_transformation, [], [f4])).
fof(f16549, plain, (~ spl10_108 | ~ spl10_118), inference(avatar_split_clause, [], [f16548, f1136, f1094])).
fof(f16548, plain, (~ (e2 = op(e0, e3)) | ~ spl10_118), inference(forward_demodulation, [], [f409, f1138])).
fof(f409, plain, ~ (op(e0, e1) = op(e0, e3)), inference(cnf_transformation, [], [f4])).
fof(f16540, plain, (~ spl10_73 | ~ spl10_48), inference(avatar_split_clause, [], [f16539, f842, f947])).
fof(f947, plain, (spl10_73 <=> (e2 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl10_73])])).
fof(f16539, plain, (~ (e2 = op(e2, e0)) | ~ spl10_48), inference(forward_demodulation, [], [f361, f844])).
fof(f361, plain, ~ (op(e2, e0) = op(e3, e0)), inference(cnf_transformation, [], [f4])).
fof(f16536, plain, (~ spl10_38 | ~ spl10_48), inference(avatar_split_clause, [], [f16535, f842, f800])).
fof(f800, plain, (spl10_38 <=> (e2 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl10_38])])).
fof(f16535, plain, (~ (e2 = op(e3, e2)) | ~ spl10_48), inference(forward_demodulation, [], [f435, f844])).
fof(f435, plain, ~ (op(e3, e0) = op(e3, e2)), inference(cnf_transformation, [], [f4])).
fof(f16527, plain, (~ spl10_33 | ~ spl10_48), inference(avatar_split_clause, [], [f16522, f842, f779])).
fof(f779, plain, (spl10_33 <=> (e2 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl10_33])])).
fof(f16522, plain, (~ (e2 = op(e3, e3)) | ~ spl10_48), inference(backward_demodulation, [], [f436, f844])).
fof(f436, plain, ~ (op(e3, e0) = op(e3, e3)), inference(cnf_transformation, [], [f4])).
fof(f16521, plain, (~ spl10_64 | ~ spl10_107 | spl10_187 | ~ spl10_228), inference(avatar_contradiction_clause, [], [f16520])).
fof(f16520, plain, ($false | (~ spl10_64 | ~ spl10_107 | spl10_187 | ~ spl10_228)), inference(subsumption_resolution, [], [f16517, f1092])).
fof(f1092, plain, ((e1 = op(e0, e3)) | ~ spl10_107), inference(avatar_component_clause, [], [f1090])).
fof(f16517, plain, (~ (e1 = op(e0, e3)) | (~ spl10_64 | spl10_187 | ~ spl10_228)), inference(backward_demodulation, [], [f16419, f911])).
fof(f911, plain, ((e3 = op(e2, e2)) | ~ spl10_64), inference(avatar_component_clause, [], [f909])).
fof(f16419, plain, (~ (e1 = op(e0, op(e2, e2))) | (spl10_187 | ~ spl10_228)), inference(forward_demodulation, [], [f1527, f1728])).
fof(f1728, plain, ((e0 = op(e2, op(e2, e2))) | ~ spl10_228), inference(avatar_component_clause, [], [f1727])).
fof(f1727, plain, (spl10_228 <=> (e0 = op(e2, op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl10_228])])).
fof(f1527, plain, (~ (e1 = op(op(e2, op(e2, e2)), op(e2, e2))) | spl10_187), inference(avatar_component_clause, [], [f1525])).
fof(f1525, plain, (spl10_187 <=> (e1 = op(op(e2, op(e2, e2)), op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl10_187])])).
fof(f16507, plain, (~ spl10_65 | ~ spl10_70), inference(avatar_split_clause, [], [f16504, f934, f913])).
fof(f913, plain, (spl10_65 <=> (e4 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl10_65])])).
fof(f16504, plain, (~ (e4 = op(e2, e2)) | ~ spl10_70), inference(backward_demodulation, [], [f428, f936])).
fof(f936, plain, ((e4 = op(e2, e1)) | ~ spl10_70), inference(avatar_component_clause, [], [f934])).
fof(f16503, plain, (~ spl10_33 | ~ spl10_83), inference(avatar_split_clause, [], [f16499, f989, f779])).
fof(f16499, plain, (~ (e2 = op(e3, e3)) | ~ spl10_83), inference(backward_demodulation, [], [f389, f991])).
fof(f389, plain, ~ (op(e1, e3) = op(e3, e3)), inference(cnf_transformation, [], [f4])).
fof(f16498, plain, (~ spl10_56 | ~ spl10_94 | ~ spl10_207 | spl10_231), inference(avatar_contradiction_clause, [], [f16497])).
fof(f16497, plain, ($false | (~ spl10_56 | ~ spl10_94 | ~ spl10_207 | spl10_231)), inference(subsumption_resolution, [], [f16491, f878])).
fof(f878, plain, ((e0 = op(e2, e3)) | ~ spl10_56), inference(avatar_component_clause, [], [f876])).
fof(f876, plain, (spl10_56 <=> (e0 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl10_56])])).
fof(f16491, plain, (~ (e0 = op(e2, e3)) | (~ spl10_94 | ~ spl10_207 | spl10_231)), inference(backward_demodulation, [], [f16417, f1037])).
fof(f1037, plain, ((e3 = op(e1, e1)) | ~ spl10_94), inference(avatar_component_clause, [], [f1035])).
fof(f16417, plain, (~ (e0 = op(e2, op(e1, e1))) | (~ spl10_207 | spl10_231)), inference(backward_demodulation, [], [f1751, f1618])).
fof(f1618, plain, ((e2 = op(e1, op(e1, e1))) | ~ spl10_207), inference(avatar_component_clause, [], [f1617])).
fof(f1617, plain, (spl10_207 <=> (e2 = op(e1, op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl10_207])])).
fof(f1751, plain, (~ (e0 = op(op(e1, op(e1, e1)), op(e1, e1))) | spl10_231), inference(avatar_component_clause, [], [f1749])).
fof(f1749, plain, (spl10_231 <=> (e0 = op(op(e1, op(e1, e1)), op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl10_231])])).
fof(f16495, plain, (spl10_35 | ~ spl10_94 | ~ spl10_242), inference(avatar_split_clause, [], [f16489, f1841, f1035, f787])).
fof(f1841, plain, (spl10_242 <=> (e4 = op(op(e1, e1), op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl10_242])])).
fof(f16489, plain, ((e4 = op(e3, e3)) | (~ spl10_94 | ~ spl10_242)), inference(backward_demodulation, [], [f1842, f1037])).
fof(f1842, plain, ((e4 = op(op(e1, e1), op(e1, e1))) | ~ spl10_242), inference(avatar_component_clause, [], [f1841])).
fof(f16494, plain, (~ spl10_33 | ~ spl10_94 | spl10_230), inference(avatar_split_clause, [], [f16486, f1745, f1035, f779])).
fof(f1745, plain, (spl10_230 <=> (e2 = op(op(e1, e1), op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl10_230])])).
fof(f16486, plain, (~ (e2 = op(e3, e3)) | (~ spl10_94 | spl10_230)), inference(backward_demodulation, [], [f1747, f1037])).
fof(f1747, plain, (~ (e2 = op(op(e1, e1), op(e1, e1))) | spl10_230), inference(avatar_component_clause, [], [f1745])).
fof(f16493, plain, (spl10_83 | ~ spl10_94 | ~ spl10_207), inference(avatar_split_clause, [], [f16485, f1617, f1035, f989])).
fof(f16485, plain, ((e2 = op(e1, e3)) | (~ spl10_94 | ~ spl10_207)), inference(backward_demodulation, [], [f1618, f1037])).
fof(f16492, plain, (~ spl10_69 | ~ spl10_94), inference(avatar_split_clause, [], [f16481, f1035, f930])).
fof(f16481, plain, (~ (e3 = op(e2, e1)) | ~ spl10_94), inference(backward_demodulation, [], [f368, f1037])).
fof(f16479, plain, (~ spl10_85 | ~ spl10_100), inference(avatar_split_clause, [], [f16475, f1060, f997])).
fof(f997, plain, (spl10_85 <=> (e4 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl10_85])])).
fof(f1060, plain, (spl10_100 <=> (e4 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl10_100])])).
fof(f16475, plain, (~ (e4 = op(e1, e3)) | ~ spl10_100), inference(backward_demodulation, [], [f416, f1062])).
fof(f1062, plain, ((e4 = op(e1, e0)) | ~ spl10_100), inference(avatar_component_clause, [], [f1060])).
fof(f416, plain, ~ (op(e1, e0) = op(e1, e3)), inference(cnf_transformation, [], [f4])).
fof(f16478, plain, (~ spl10_50 | ~ spl10_100), inference(avatar_split_clause, [], [f16473, f1060, f850])).
fof(f850, plain, (spl10_50 <=> (e4 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl10_50])])).
fof(f16473, plain, (~ (e4 = op(e3, e0)) | ~ spl10_100), inference(backward_demodulation, [], [f359, f1062])).
fof(f359, plain, ~ (op(e1, e0) = op(e3, e0)), inference(cnf_transformation, [], [f4])).
fof(f16471, plain, (~ spl10_65 | ~ spl10_115), inference(avatar_split_clause, [], [f16468, f1123, f913])).
fof(f1123, plain, (spl10_115 <=> (e4 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl10_115])])).
fof(f16468, plain, (~ (e4 = op(e2, e2)) | ~ spl10_115), inference(backward_demodulation, [], [f375, f1125])).
fof(f1125, plain, ((e4 = op(e0, e2)) | ~ spl10_115), inference(avatar_component_clause, [], [f1123])).
fof(f16467, plain, (~ spl10_93 | ~ spl10_118), inference(avatar_split_clause, [], [f16460, f1136, f1031])).
fof(f1031, plain, (spl10_93 <=> (e2 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl10_93])])).
fof(f16460, plain, (~ (e2 = op(e1, e1)) | ~ spl10_118), inference(backward_demodulation, [], [f364, f1138])).
fof(f364, plain, ~ (op(e0, e1) = op(e1, e1)), inference(cnf_transformation, [], [f4])).
fof(f16459, plain, (~ spl10_83 | ~ spl10_124 | spl10_221 | ~ spl10_236), inference(avatar_split_clause, [], [f16454, f1781, f1691, f1161, f989])).
fof(f1161, plain, (spl10_124 <=> (op(e0, e0) = e3)), introduced(avatar_definition, [new_symbols(naming, [spl10_124])])).
fof(f1691, plain, (spl10_221 <=> (e2 = op(op(e0, op(e0, e0)), op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl10_221])])).
fof(f1781, plain, (spl10_236 <=> (e1 = op(e0, op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl10_236])])).
fof(f16454, plain, (~ (e2 = op(e1, e3)) | (~ spl10_124 | spl10_221 | ~ spl10_236)), inference(backward_demodulation, [], [f16413, f1163])).
fof(f1163, plain, ((op(e0, e0) = e3) | ~ spl10_124), inference(avatar_component_clause, [], [f1161])).
fof(f16413, plain, (~ (e2 = op(e1, op(e0, e0))) | (spl10_221 | ~ spl10_236)), inference(forward_demodulation, [], [f1693, f1782])).
fof(f1782, plain, ((e1 = op(e0, op(e0, e0))) | ~ spl10_236), inference(avatar_component_clause, [], [f1781])).
fof(f1693, plain, (~ (e2 = op(op(e0, op(e0, e0)), op(e0, e0))) | spl10_221), inference(avatar_component_clause, [], [f1691])).
fof(f16457, plain, (~ spl10_33 | ~ spl10_124 | spl10_234), inference(avatar_split_clause, [], [f16449, f1767, f1161, f779])).
fof(f1767, plain, (spl10_234 <=> (e2 = op(op(e0, e0), op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl10_234])])).
fof(f16449, plain, (~ (e2 = op(e3, e3)) | (~ spl10_124 | spl10_234)), inference(backward_demodulation, [], [f1769, f1163])).
fof(f1769, plain, (~ (e2 = op(op(e0, e0), op(e0, e0))) | spl10_234), inference(avatar_component_clause, [], [f1767])).
fof(f16456, plain, (~ spl10_114 | ~ spl10_124), inference(avatar_split_clause, [], [f16446, f1161, f1119])).
fof(f16446, plain, (~ (e3 = op(e0, e2)) | ~ spl10_124), inference(backward_demodulation, [], [f405, f1163])).
fof(f405, plain, ~ (op(e0, e0) = op(e0, e2)), inference(cnf_transformation, [], [f4])).
fof(f16455, plain, (~ spl10_99 | ~ spl10_124), inference(avatar_split_clause, [], [f16443, f1161, f1056])).
fof(f16443, plain, (~ (e3 = op(e1, e0)) | ~ spl10_124), inference(backward_demodulation, [], [f354, f1163])).
fof(f354, plain, ~ (op(e0, e0) = op(e1, e0)), inference(cnf_transformation, [], [f4])).
fof(f16441, plain, (~ spl10_65 | ~ spl10_53 | spl10_162), inference(avatar_split_clause, [], [f16440, f1386, f863, f913])).
fof(f863, plain, (spl10_53 <=> (e2 = op(e2, e4))), introduced(avatar_definition, [new_symbols(naming, [spl10_53])])).
fof(f1386, plain, (spl10_162 <=> (e4 = op(e2, op(e2, e4)))), introduced(avatar_definition, [new_symbols(naming, [spl10_162])])).
fof(f16440, plain, (~ (e4 = op(e2, e2)) | (~ spl10_53 | spl10_162)), inference(forward_demodulation, [], [f1388, f865])).
fof(f865, plain, ((e2 = op(e2, e4)) | ~ spl10_53), inference(avatar_component_clause, [], [f863])).
fof(f1388, plain, (~ (e4 = op(e2, op(e2, e4))) | spl10_162), inference(avatar_component_clause, [], [f1386])).
fof(f16437, plain, (spl10_48 | ~ spl10_56 | ~ spl10_164), inference(avatar_split_clause, [], [f16436, f1401, f876, f842])).
fof(f1401, plain, (spl10_164 <=> (e2 = op(e3, op(e2, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl10_164])])).
fof(f16436, plain, ((e2 = op(e3, e0)) | (~ spl10_56 | ~ spl10_164)), inference(forward_demodulation, [], [f1403, f878])).
fof(f1403, plain, ((e2 = op(e3, op(e2, e3))) | ~ spl10_164), inference(avatar_component_clause, [], [f1401])).
fof(f16422, plain, (spl10_118 | ~ spl10_72 | ~ spl10_180), inference(avatar_split_clause, [], [f16421, f1489, f943, f1136])).
fof(f1489, plain, (spl10_180 <=> (e2 = op(e0, op(e2, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl10_180])])).
fof(f16421, plain, ((e2 = op(e0, e1)) | (~ spl10_72 | ~ spl10_180)), inference(forward_demodulation, [], [f1491, f945])).
fof(f945, plain, ((e1 = op(e2, e0)) | ~ spl10_72), inference(avatar_component_clause, [], [f943])).
fof(f1491, plain, ((e2 = op(e0, op(e2, e0))) | ~ spl10_180), inference(avatar_component_clause, [], [f1489])).
fof(f16400, plain, (~ spl10_125 | ~ spl10_101 | spl10_136), inference(avatar_split_clause, [], [f16206, f1264, f1065, f1165])).
fof(f1165, plain, (spl10_125 <=> (op(e0, e0) = e4)), introduced(avatar_definition, [new_symbols(naming, [spl10_125])])).
fof(f1065, plain, (spl10_101 <=> (e0 = op(e0, e4))), introduced(avatar_definition, [new_symbols(naming, [spl10_101])])).
fof(f1264, plain, (spl10_136 <=> (e4 = op(e0, op(e0, e4)))), introduced(avatar_definition, [new_symbols(naming, [spl10_136])])).
fof(f16206, plain, (~ (op(e0, e0) = e4) | (~ spl10_101 | spl10_136)), inference(backward_demodulation, [], [f1266, f1067])).
fof(f1067, plain, ((e0 = op(e0, e4)) | ~ spl10_101), inference(avatar_component_clause, [], [f1065])).
fof(f1266, plain, (~ (e4 = op(e0, op(e0, e4))) | spl10_136), inference(avatar_component_clause, [], [f1264])).
fof(f16399, plain, (~ spl10_125 | ~ spl10_21 | spl10_182), inference(avatar_split_clause, [], [f16330, f1497, f729, f1165])).
fof(f729, plain, (spl10_21 <=> (e0 = op(e4, e0))), introduced(avatar_definition, [new_symbols(naming, [spl10_21])])).
fof(f1497, plain, (spl10_182 <=> (e4 = op(e0, op(e4, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl10_182])])).
fof(f16330, plain, (~ (op(e0, e0) = e4) | (~ spl10_21 | spl10_182)), inference(backward_demodulation, [], [f1498, f731])).
fof(f731, plain, ((e0 = op(e4, e0)) | ~ spl10_21), inference(avatar_component_clause, [], [f729])).
fof(f1498, plain, (~ (e4 = op(e0, op(e4, e0))) | spl10_182), inference(avatar_component_clause, [], [f1497])).
fof(f16332, plain, (~ spl10_6 | ~ spl10_21), inference(avatar_split_clause, [], [f16324, f729, f666])).
fof(f666, plain, (spl10_6 <=> (e0 = op(e4, e3))), introduced(avatar_definition, [new_symbols(naming, [spl10_6])])).
fof(f16324, plain, (~ (e0 = op(e4, e3)) | ~ spl10_21), inference(backward_demodulation, [], [f446, f731])).
fof(f446, plain, ~ (op(e4, e0) = op(e4, e3)), inference(cnf_transformation, [], [f4])).
fof(f16315, plain, (~ spl10_33 | ~ spl10_86 | ~ spl10_203 | spl10_217), inference(avatar_contradiction_clause, [], [f16314])).
fof(f16314, plain, ($false | (~ spl10_33 | ~ spl10_86 | ~ spl10_203 | spl10_217)), inference(subsumption_resolution, [], [f16310, f1004])).
fof(f1004, plain, ((e0 = op(e1, e2)) | ~ spl10_86), inference(avatar_component_clause, [], [f1002])).
fof(f1002, plain, (spl10_86 <=> (e0 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl10_86])])).
fof(f16310, plain, (~ (e0 = op(e1, e2)) | (~ spl10_33 | ~ spl10_203 | spl10_217)), inference(backward_demodulation, [], [f16124, f781])).
fof(f781, plain, ((e2 = op(e3, e3)) | ~ spl10_33), inference(avatar_component_clause, [], [f779])).
fof(f16124, plain, (~ (e0 = op(e1, op(e3, e3))) | (~ spl10_203 | spl10_217)), inference(backward_demodulation, [], [f1673, f1598])).
fof(f1598, plain, ((e1 = op(e3, op(e3, e3))) | ~ spl10_203), inference(avatar_component_clause, [], [f1597])).
fof(f1597, plain, (spl10_203 <=> (e1 = op(e3, op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl10_203])])).
fof(f1673, plain, (~ (e0 = op(op(e3, op(e3, e3)), op(e3, e3))) | spl10_217), inference(avatar_component_clause, [], [f1671])).
fof(f1671, plain, (spl10_217 <=> (e0 = op(op(e3, op(e3, e3)), op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl10_217])])).
fof(f16311, plain, (~ spl10_28 | ~ spl10_33), inference(avatar_split_clause, [], [f16300, f779, f758])).
fof(f758, plain, (spl10_28 <=> (e2 = op(e3, e4))), introduced(avatar_definition, [new_symbols(naming, [spl10_28])])).
fof(f16300, plain, (~ (e2 = op(e3, e4)) | ~ spl10_33), inference(backward_demodulation, [], [f443, f781])).
fof(f443, plain, ~ (op(e3, e3) = op(e3, e4)), inference(cnf_transformation, [], [f4])).
fof(f16297, plain, (~ spl10_32 | ~ spl10_37), inference(avatar_split_clause, [], [f16291, f796, f775])).
fof(f16291, plain, (~ (e1 = op(e3, e3)) | ~ spl10_37), inference(backward_demodulation, [], [f441, f798])).
fof(f798, plain, ((e1 = op(e3, e2)) | ~ spl10_37), inference(avatar_component_clause, [], [f796])).
fof(f16279, plain, (~ spl10_6 | ~ spl10_56), inference(avatar_split_clause, [], [f16271, f876, f666])).
fof(f16271, plain, (~ (e0 = op(e4, e3)) | ~ spl10_56), inference(backward_demodulation, [], [f392, f878])).
fof(f392, plain, ~ (op(e2, e3) = op(e4, e3)), inference(cnf_transformation, [], [f4])).
fof(f16269, plain, (~ spl10_2 | ~ spl10_65 | spl10_214), inference(avatar_split_clause, [], [f16261, f1657, f913, f649])).
fof(f649, plain, (spl10_2 <=> (e1 = op(e4, e4))), introduced(avatar_definition, [new_symbols(naming, [spl10_2])])).
fof(f1657, plain, (spl10_214 <=> (e1 = op(op(e2, e2), op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl10_214])])).
fof(f16261, plain, (~ (e1 = op(e4, e4)) | (~ spl10_65 | spl10_214)), inference(backward_demodulation, [], [f1659, f915])).
fof(f915, plain, ((e4 = op(e2, e2)) | ~ spl10_65), inference(avatar_component_clause, [], [f913])).
fof(f1659, plain, (~ (e1 = op(op(e2, e2), op(e2, e2))) | spl10_214), inference(avatar_component_clause, [], [f1657])).
fof(f16267, plain, (~ spl10_53 | ~ spl10_65 | spl10_169), inference(avatar_split_clause, [], [f16256, f1429, f913, f863])).
fof(f1429, plain, (spl10_169 <=> (e2 = op(e2, op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl10_169])])).
fof(f16256, plain, (~ (e2 = op(e2, e4)) | (~ spl10_65 | spl10_169)), inference(backward_demodulation, [], [f1430, f915])).
fof(f1430, plain, (~ (e2 = op(e2, op(e2, e2))) | spl10_169), inference(avatar_component_clause, [], [f1429])).
fof(f16265, plain, (~ spl10_15 | ~ spl10_65), inference(avatar_split_clause, [], [f16253, f913, f703])).
fof(f703, plain, (spl10_15 <=> (e4 = op(e4, e2))), introduced(avatar_definition, [new_symbols(naming, [spl10_15])])).
fof(f16253, plain, (~ (e4 = op(e4, e2)) | ~ spl10_65), inference(backward_demodulation, [], [f382, f915])).
fof(f382, plain, ~ (op(e2, e2) = op(e4, e2)), inference(cnf_transformation, [], [f4])).
fof(f16251, plain, (~ spl10_59 | ~ spl10_69), inference(avatar_split_clause, [], [f16245, f930, f888])).
fof(f888, plain, (spl10_59 <=> (e3 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl10_59])])).
fof(f16245, plain, (~ (e3 = op(e2, e3)) | ~ spl10_69), inference(backward_demodulation, [], [f429, f932])).
fof(f429, plain, ~ (op(e2, e1) = op(e2, e3)), inference(cnf_transformation, [], [f4])).
fof(f16250, plain, (~ spl10_19 | ~ spl10_69), inference(avatar_split_clause, [], [f16243, f930, f720])).
fof(f720, plain, (spl10_19 <=> (e3 = op(e4, e1))), introduced(avatar_definition, [new_symbols(naming, [spl10_19])])).
fof(f16243, plain, (~ (e3 = op(e4, e1)) | ~ spl10_69), inference(backward_demodulation, [], [f372, f932])).
fof(f372, plain, ~ (op(e2, e1) = op(e4, e1)), inference(cnf_transformation, [], [f4])).
fof(f16242, plain, (~ spl10_67 | ~ spl10_72), inference(avatar_split_clause, [], [f16234, f943, f922])).
fof(f922, plain, (spl10_67 <=> (e1 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl10_67])])).
fof(f16234, plain, (~ (e1 = op(e2, e1)) | ~ spl10_72), inference(backward_demodulation, [], [f424, f945])).
fof(f16223, plain, (spl10_72 | ~ spl10_86 | ~ spl10_144), inference(avatar_split_clause, [], [f16218, f1301, f1002, f943])).
fof(f1301, plain, (spl10_144 <=> (e1 = op(e2, op(e1, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl10_144])])).
fof(f16218, plain, ((e1 = op(e2, e0)) | (~ spl10_86 | ~ spl10_144)), inference(backward_demodulation, [], [f1303, f1004])).
fof(f1303, plain, ((e1 = op(e2, op(e1, e2))) | ~ spl10_144), inference(avatar_component_clause, [], [f1301])).
fof(f16222, plain, (~ spl10_76 | ~ spl10_86), inference(avatar_split_clause, [], [f16213, f1002, f960])).
fof(f960, plain, (spl10_76 <=> (e0 = op(e1, e4))), introduced(avatar_definition, [new_symbols(naming, [spl10_76])])).
fof(f16213, plain, (~ (e0 = op(e1, e4)) | ~ spl10_86), inference(backward_demodulation, [], [f422, f1004])).
fof(f422, plain, ~ (op(e1, e2) = op(e1, e4)), inference(cnf_transformation, [], [f4])).
fof(f16207, plain, (~ spl10_76 | ~ spl10_101), inference(avatar_split_clause, [], [f16197, f1065, f960])).
fof(f16197, plain, (~ (e0 = op(e1, e4)) | ~ spl10_101), inference(backward_demodulation, [], [f394, f1067])).
fof(f394, plain, ~ (op(e0, e4) = op(e1, e4)), inference(cnf_transformation, [], [f4])).
fof(f16196, plain, (~ spl10_32 | ~ spl10_107), inference(avatar_split_clause, [], [f16188, f1090, f775])).
fof(f16188, plain, (~ (e1 = op(e3, e3)) | ~ spl10_107), inference(backward_demodulation, [], [f386, f1092])).
fof(f16186, plain, (~ spl10_108 | ~ spl10_114 | spl10_148), inference(avatar_split_clause, [], [f16182, f1320, f1119, f1094])).
fof(f1320, plain, (spl10_148 <=> (e2 = op(e0, op(e0, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl10_148])])).
fof(f16182, plain, (~ (e2 = op(e0, e3)) | (~ spl10_114 | spl10_148)), inference(backward_demodulation, [], [f1322, f1121])).
fof(f1322, plain, (~ (e2 = op(e0, op(e0, e2))) | spl10_148), inference(avatar_component_clause, [], [f1320])).
fof(f16185, plain, (spl10_56 | ~ spl10_114 | ~ spl10_147), inference(avatar_split_clause, [], [f16181, f1315, f1119, f876])).
fof(f1315, plain, (spl10_147 <=> (e0 = op(e2, op(e0, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl10_147])])).
fof(f16181, plain, ((e0 = op(e2, e3)) | (~ spl10_114 | ~ spl10_147)), inference(backward_demodulation, [], [f1317, f1121])).
fof(f1317, plain, ((e0 = op(e2, op(e0, e2))) | ~ spl10_147), inference(avatar_component_clause, [], [f1315])).
fof(f16184, plain, (~ spl10_104 | ~ spl10_114), inference(avatar_split_clause, [], [f16179, f1119, f1077])).
fof(f1077, plain, (spl10_104 <=> (e3 = op(e0, e4))), introduced(avatar_definition, [new_symbols(naming, [spl10_104])])).
fof(f16179, plain, (~ (e3 = op(e0, e4)) | ~ spl10_114), inference(backward_demodulation, [], [f412, f1121])).
fof(f412, plain, ~ (op(e0, e2) = op(e0, e4)), inference(cnf_transformation, [], [f4])).
fof(f16183, plain, (~ spl10_39 | ~ spl10_114), inference(avatar_split_clause, [], [f16176, f1119, f804])).
fof(f804, plain, (spl10_39 <=> (e3 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl10_39])])).
fof(f16176, plain, (~ (e3 = op(e3, e2)) | ~ spl10_114), inference(backward_demodulation, [], [f376, f1121])).
fof(f376, plain, ~ (op(e0, e2) = op(e3, e2)), inference(cnf_transformation, [], [f4])).
fof(f16173, plain, (~ spl10_37 | ~ spl10_123 | ~ spl10_222 | spl10_235), inference(avatar_split_clause, [], [f16166, f1771, f1695, f1157, f796])).
fof(f1157, plain, (spl10_123 <=> (op(e0, e0) = e2)), introduced(avatar_definition, [new_symbols(naming, [spl10_123])])).
fof(f1695, plain, (spl10_222 <=> (e3 = op(e0, op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl10_222])])).
fof(f1771, plain, (spl10_235 <=> (e1 = op(op(e0, op(e0, e0)), op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl10_235])])).
fof(f16166, plain, (~ (e1 = op(e3, e2)) | (~ spl10_123 | ~ spl10_222 | spl10_235)), inference(backward_demodulation, [], [f16118, f1159])).
fof(f1159, plain, ((op(e0, e0) = e2) | ~ spl10_123), inference(avatar_component_clause, [], [f1157])).
fof(f16118, plain, (~ (e1 = op(e3, op(e0, e0))) | (~ spl10_222 | spl10_235)), inference(backward_demodulation, [], [f1773, f1696])).
fof(f1696, plain, ((e3 = op(e0, op(e0, e0))) | ~ spl10_222), inference(avatar_component_clause, [], [f1695])).
fof(f1773, plain, (~ (e1 = op(op(e0, op(e0, e0)), op(e0, e0))) | spl10_235), inference(avatar_component_clause, [], [f1771])).
fof(f16171, plain, (spl10_65 | ~ spl10_123 | ~ spl10_245), inference(avatar_split_clause, [], [f16164, f1859, f1157, f913])).
fof(f1859, plain, (spl10_245 <=> (e4 = op(op(e0, e0), op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl10_245])])).
fof(f16164, plain, ((e4 = op(e2, e2)) | (~ spl10_123 | ~ spl10_245)), inference(backward_demodulation, [], [f1860, f1159])).
fof(f1860, plain, ((e4 = op(op(e0, e0), op(e0, e0))) | ~ spl10_245), inference(avatar_component_clause, [], [f1859])).
fof(f16169, plain, (spl10_114 | ~ spl10_123 | ~ spl10_222), inference(avatar_split_clause, [], [f16159, f1695, f1157, f1119])).
fof(f16159, plain, ((e3 = op(e0, e2)) | (~ spl10_123 | ~ spl10_222)), inference(backward_demodulation, [], [f1696, f1159])).
fof(f16168, plain, (~ spl10_108 | ~ spl10_123), inference(avatar_split_clause, [], [f16154, f1157, f1094])).
fof(f16154, plain, (~ (e2 = op(e0, e3)) | ~ spl10_123), inference(backward_demodulation, [], [f406, f1159])).
fof(f16167, plain, (~ spl10_23 | ~ spl10_123), inference(avatar_split_clause, [], [f16152, f1157, f737])).
fof(f737, plain, (spl10_23 <=> (e2 = op(e4, e0))), introduced(avatar_definition, [new_symbols(naming, [spl10_23])])).
fof(f16152, plain, (~ (e2 = op(e4, e0)) | ~ spl10_123), inference(backward_demodulation, [], [f357, f1159])).
fof(f357, plain, ~ (op(e0, e0) = op(e4, e0)), inference(cnf_transformation, [], [f4])).
fof(f16150, plain, (spl10_5 | ~ spl10_130), inference(avatar_split_clause, [], [f16140, f1186, f661])).
fof(f661, plain, (spl10_5 <=> (e4 = op(e4, e4))), introduced(avatar_definition, [new_symbols(naming, [spl10_5])])).
fof(f1186, plain, (spl10_130 <=> (e4 = unit)), introduced(avatar_definition, [new_symbols(naming, [spl10_130])])).
fof(f16140, plain, ((e4 = op(e4, e4)) | ~ spl10_130), inference(backward_demodulation, [], [f302, f1188])).
fof(f1188, plain, ((e4 = unit) | ~ spl10_130), inference(avatar_component_clause, [], [f1186])).
fof(f302, plain, (e4 = op(e4, unit)), inference(cnf_transformation, [], [f2])).
fof(f2, plain, (((e4 = unit) | (e3 = unit) | (e2 = unit) | (e1 = unit) | (e0 = unit)) & (e4 = op(e4, unit)) & (e4 = op(unit, e4)) & (e3 = op(e3, unit)) & (e3 = op(unit, e3)) & (e2 = op(e2, unit)) & (e2 = op(unit, e2)) & (e1 = op(e1, unit)) & (e1 = op(unit, e1)) & (e0 = op(e0, unit)) & (e0 = op(unit, e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax2)).
fof(f16148, plain, (spl10_29 | ~ spl10_130), inference(avatar_split_clause, [], [f16138, f1186, f762])).
fof(f16138, plain, ((e3 = op(e3, e4)) | ~ spl10_130), inference(backward_demodulation, [], [f300, f1188])).
fof(f300, plain, (e3 = op(e3, unit)), inference(cnf_transformation, [], [f2])).
fof(f16147, plain, (spl10_9 | ~ spl10_130), inference(avatar_split_clause, [], [f16137, f1186, f678])).
fof(f678, plain, (spl10_9 <=> (e3 = op(e4, e3))), introduced(avatar_definition, [new_symbols(naming, [spl10_9])])).
fof(f16137, plain, ((e3 = op(e4, e3)) | ~ spl10_130), inference(backward_demodulation, [], [f299, f1188])).
fof(f299, plain, (e3 = op(unit, e3)), inference(cnf_transformation, [], [f2])).
fof(f16146, plain, (spl10_53 | ~ spl10_130), inference(avatar_split_clause, [], [f16136, f1186, f863])).
fof(f16136, plain, ((e2 = op(e2, e4)) | ~ spl10_130), inference(backward_demodulation, [], [f298, f1188])).
fof(f298, plain, (e2 = op(e2, unit)), inference(cnf_transformation, [], [f2])).
fof(f16145, plain, (spl10_13 | ~ spl10_130), inference(avatar_split_clause, [], [f16135, f1186, f695])).
fof(f695, plain, (spl10_13 <=> (e2 = op(e4, e2))), introduced(avatar_definition, [new_symbols(naming, [spl10_13])])).
fof(f16135, plain, ((e2 = op(e4, e2)) | ~ spl10_130), inference(backward_demodulation, [], [f297, f1188])).
fof(f297, plain, (e2 = op(unit, e2)), inference(cnf_transformation, [], [f2])).
fof(f16144, plain, (spl10_77 | ~ spl10_130), inference(avatar_split_clause, [], [f16134, f1186, f964])).
fof(f964, plain, (spl10_77 <=> (e1 = op(e1, e4))), introduced(avatar_definition, [new_symbols(naming, [spl10_77])])).
fof(f16134, plain, ((e1 = op(e1, e4)) | ~ spl10_130), inference(backward_demodulation, [], [f296, f1188])).
fof(f296, plain, (e1 = op(e1, unit)), inference(cnf_transformation, [], [f2])).
fof(f16143, plain, (spl10_17 | ~ spl10_130), inference(avatar_split_clause, [], [f16133, f1186, f712])).
fof(f712, plain, (spl10_17 <=> (e1 = op(e4, e1))), introduced(avatar_definition, [new_symbols(naming, [spl10_17])])).
fof(f16133, plain, ((e1 = op(e4, e1)) | ~ spl10_130), inference(backward_demodulation, [], [f295, f1188])).
fof(f295, plain, (e1 = op(unit, e1)), inference(cnf_transformation, [], [f2])).
fof(f16142, plain, (spl10_101 | ~ spl10_130), inference(avatar_split_clause, [], [f16132, f1186, f1065])).
fof(f16132, plain, ((e0 = op(e0, e4)) | ~ spl10_130), inference(backward_demodulation, [], [f294, f1188])).
fof(f294, plain, (e0 = op(e0, unit)), inference(cnf_transformation, [], [f2])).
fof(f16141, plain, (spl10_21 | ~ spl10_130), inference(avatar_split_clause, [], [f16131, f1186, f729])).
fof(f16131, plain, ((e0 = op(e4, e0)) | ~ spl10_130), inference(backward_demodulation, [], [f293, f1188])).
fof(f293, plain, (e0 = op(unit, e0)), inference(cnf_transformation, [], [f2])).
fof(f16130, plain, (~ spl10_76 | ~ spl10_120 | spl10_152), inference(avatar_split_clause, [], [f16129, f1339, f1144, f960])).
fof(f1144, plain, (spl10_120 <=> (e4 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl10_120])])).
fof(f1339, plain, (spl10_152 <=> (e0 = op(e1, op(e0, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl10_152])])).
fof(f16129, plain, (~ (e0 = op(e1, e4)) | (~ spl10_120 | spl10_152)), inference(forward_demodulation, [], [f1340, f1146])).
fof(f1146, plain, ((e4 = op(e0, e1)) | ~ spl10_120), inference(avatar_component_clause, [], [f1144])).
fof(f1340, plain, (~ (e0 = op(e1, op(e0, e1))) | spl10_152), inference(avatar_component_clause, [], [f1339])).
fof(f16128, plain, (spl10_107 | ~ spl10_99 | ~ spl10_155), inference(avatar_split_clause, [], [f16127, f1353, f1056, f1090])).
fof(f1353, plain, (spl10_155 <=> (e1 = op(e0, op(e1, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl10_155])])).
fof(f16127, plain, ((e1 = op(e0, e3)) | (~ spl10_99 | ~ spl10_155)), inference(forward_demodulation, [], [f1355, f1058])).
fof(f1355, plain, ((e1 = op(e0, op(e1, e0))) | ~ spl10_155), inference(avatar_component_clause, [], [f1353])).
fof(f16126, plain, (~ spl10_104 | ~ spl10_50 | spl10_181), inference(avatar_split_clause, [], [f16125, f1493, f850, f1077])).
fof(f16125, plain, (~ (e3 = op(e0, e4)) | (~ spl10_50 | spl10_181)), inference(forward_demodulation, [], [f1494, f852])).
fof(f852, plain, ((e4 = op(e3, e0)) | ~ spl10_50), inference(avatar_component_clause, [], [f850])).
fof(f1494, plain, (~ (e3 = op(e0, op(e3, e0))) | spl10_181), inference(avatar_component_clause, [], [f1493])).
fof(f16113, plain, (spl10_65 | ~ spl10_93 | ~ spl10_242), inference(avatar_split_clause, [], [f16112, f1841, f1031, f913])).
fof(f16112, plain, ((e4 = op(e2, e2)) | (~ spl10_93 | ~ spl10_242)), inference(forward_demodulation, [], [f1842, f1033])).
fof(f1033, plain, ((e2 = op(e1, e1)) | ~ spl10_93), inference(avatar_component_clause, [], [f1031])).
fof(f16108, plain, (~ spl10_110 | ~ spl10_85), inference(avatar_split_clause, [], [f16107, f997, f1102])).
fof(f16107, plain, (~ (e4 = op(e0, e3)) | ~ spl10_85), inference(forward_demodulation, [], [f384, f999])).
fof(f999, plain, ((e4 = op(e1, e3)) | ~ spl10_85), inference(avatar_component_clause, [], [f997])).
fof(f16105, plain, (~ spl10_86 | ~ spl10_93 | spl10_237), inference(avatar_split_clause, [], [f15921, f1787, f1031, f1002])).
fof(f1787, plain, (spl10_237 <=> (e0 = op(e1, op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl10_237])])).
fof(f15921, plain, (~ (e0 = op(e1, e2)) | (~ spl10_93 | spl10_237)), inference(backward_demodulation, [], [f1789, f1033])).
fof(f1789, plain, (~ (e0 = op(e1, op(e1, e1))) | spl10_237), inference(avatar_component_clause, [], [f1787])).
fof(f16098, plain, (~ spl10_2 | ~ spl10_93 | spl10_233), inference(avatar_contradiction_clause, [], [f16097])).
fof(f16097, plain, ($false | (~ spl10_2 | ~ spl10_93 | spl10_233)), inference(subsumption_resolution, [], [f16094, f1033])).
fof(f16094, plain, (~ (e2 = op(e1, e1)) | (~ spl10_2 | spl10_233)), inference(backward_demodulation, [], [f1763, f651])).
fof(f651, plain, ((e1 = op(e4, e4)) | ~ spl10_2), inference(avatar_component_clause, [], [f649])).
fof(f1763, plain, (~ (e2 = op(op(e4, e4), op(e4, e4))) | spl10_233), inference(avatar_component_clause, [], [f1761])).
fof(f1761, plain, (spl10_233 <=> (e2 = op(op(e4, e4), op(e4, e4)))), introduced(avatar_definition, [new_symbols(naming, [spl10_233])])).
fof(f16077, plain, (~ spl10_5 | ~ spl10_15), inference(avatar_split_clause, [], [f16074, f703, f661])).
fof(f16074, plain, (~ (e4 = op(e4, e4)) | ~ spl10_15), inference(backward_demodulation, [], [f452, f705])).
fof(f705, plain, ((e4 = op(e4, e2)) | ~ spl10_15), inference(avatar_component_clause, [], [f703])).
fof(f452, plain, ~ (op(e4, e2) = op(e4, e4)), inference(cnf_transformation, [], [f4])).
fof(f16072, plain, (spl10_7 | ~ spl10_19 | ~ spl10_179), inference(avatar_contradiction_clause, [], [f16071])).
fof(f16071, plain, ($false | (spl10_7 | ~ spl10_19 | ~ spl10_179)), inference(subsumption_resolution, [], [f16069, f671])).
fof(f671, plain, (~ (e1 = op(e4, e3)) | spl10_7), inference(avatar_component_clause, [], [f670])).
fof(f670, plain, (spl10_7 <=> (e1 = op(e4, e3))), introduced(avatar_definition, [new_symbols(naming, [spl10_7])])).
fof(f16069, plain, ((e1 = op(e4, e3)) | (~ spl10_19 | ~ spl10_179)), inference(backward_demodulation, [], [f1482, f722])).
fof(f722, plain, ((e3 = op(e4, e1)) | ~ spl10_19), inference(avatar_component_clause, [], [f720])).
fof(f1482, plain, ((e1 = op(e4, op(e4, e1))) | ~ spl10_179), inference(avatar_component_clause, [], [f1481])).
fof(f1481, plain, (spl10_179 <=> (e1 = op(e4, op(e4, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl10_179])])).
fof(f16070, plain, (~ spl10_9 | ~ spl10_19), inference(avatar_split_clause, [], [f16065, f720, f678])).
fof(f16065, plain, (~ (e3 = op(e4, e3)) | ~ spl10_19), inference(backward_demodulation, [], [f449, f722])).
fof(f449, plain, ~ (op(e4, e1) = op(e4, e3)), inference(cnf_transformation, [], [f4])).
fof(f16062, plain, (~ spl10_13 | ~ spl10_23), inference(avatar_split_clause, [], [f16055, f737, f695])).
fof(f16055, plain, (~ (e2 = op(e4, e2)) | ~ spl10_23), inference(backward_demodulation, [], [f445, f739])).
fof(f739, plain, ((e2 = op(e4, e0)) | ~ spl10_23), inference(avatar_component_clause, [], [f737])).
fof(f445, plain, ~ (op(e4, e0) = op(e4, e2)), inference(cnf_transformation, [], [f4])).
fof(f16047, plain, (~ spl10_32 | ~ spl10_120 | spl10_213 | ~ spl10_224), inference(avatar_contradiction_clause, [], [f16046])).
fof(f16046, plain, ($false | (~ spl10_32 | ~ spl10_120 | spl10_213 | ~ spl10_224)), inference(subsumption_resolution, [], [f16043, f1146])).
fof(f16043, plain, (~ (e4 = op(e0, e1)) | (~ spl10_32 | spl10_213 | ~ spl10_224)), inference(backward_demodulation, [], [f15803, f777])).
fof(f15803, plain, (~ (e4 = op(e0, op(e3, e3))) | (spl10_213 | ~ spl10_224)), inference(backward_demodulation, [], [f1653, f1706])).
fof(f1653, plain, (~ (e4 = op(op(e3, op(e3, e3)), op(e3, e3))) | spl10_213), inference(avatar_component_clause, [], [f1651])).
fof(f1651, plain, (spl10_213 <=> (e4 = op(op(e3, op(e3, e3)), op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl10_213])])).
fof(f16034, plain, (~ spl10_29 | ~ spl10_39), inference(avatar_split_clause, [], [f16030, f804, f762])).
fof(f16030, plain, (~ (e3 = op(e3, e4)) | ~ spl10_39), inference(backward_demodulation, [], [f442, f806])).
fof(f806, plain, ((e3 = op(e3, e2)) | ~ spl10_39), inference(avatar_component_clause, [], [f804])).
fof(f442, plain, ~ (op(e3, e2) = op(e3, e4)), inference(cnf_transformation, [], [f4])).
fof(f16027, plain, (~ spl10_36 | ~ spl10_41), inference(avatar_split_clause, [], [f16018, f813, f792])).
fof(f792, plain, (spl10_36 <=> (e0 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl10_36])])).
fof(f16018, plain, (~ (e0 = op(e3, e2)) | ~ spl10_41), inference(backward_demodulation, [], [f438, f815])).
fof(f438, plain, ~ (op(e3, e1) = op(e3, e2)), inference(cnf_transformation, [], [f4])).
fof(f16016, plain, (~ spl10_35 | ~ spl10_50), inference(avatar_split_clause, [], [f16012, f850, f787])).
fof(f16012, plain, (~ (e4 = op(e3, e3)) | ~ spl10_50), inference(backward_demodulation, [], [f436, f852])).
fof(f16008, plain, (~ spl10_5 | ~ spl10_55), inference(avatar_split_clause, [], [f16005, f871, f661])).
fof(f871, plain, (spl10_55 <=> (e4 = op(e2, e4))), introduced(avatar_definition, [new_symbols(naming, [spl10_55])])).
fof(f16005, plain, (~ (e4 = op(e4, e4)) | ~ spl10_55), inference(backward_demodulation, [], [f402, f873])).
fof(f873, plain, ((e4 = op(e2, e4)) | ~ spl10_55), inference(avatar_component_clause, [], [f871])).
fof(f402, plain, ~ (op(e2, e4) = op(e4, e4)), inference(cnf_transformation, [], [f4])).
fof(f16003, plain, (~ spl10_9 | ~ spl10_59), inference(avatar_split_clause, [], [f15998, f888, f678])).
fof(f15998, plain, (~ (e3 = op(e4, e3)) | ~ spl10_59), inference(backward_demodulation, [], [f392, f890])).
fof(f890, plain, ((e3 = op(e2, e3)) | ~ spl10_59), inference(avatar_component_clause, [], [f888])).
fof(f15996, plain, (~ spl10_53 | ~ spl10_63), inference(avatar_split_clause, [], [f15988, f905, f863])).
fof(f905, plain, (spl10_63 <=> (e2 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl10_63])])).
fof(f15988, plain, (~ (e2 = op(e2, e4)) | ~ spl10_63), inference(backward_demodulation, [], [f432, f907])).
fof(f907, plain, ((e2 = op(e2, e2)) | ~ spl10_63), inference(avatar_component_clause, [], [f905])).
fof(f432, plain, ~ (op(e2, e2) = op(e2, e4)), inference(cnf_transformation, [], [f4])).
fof(f15995, plain, (~ spl10_13 | ~ spl10_63), inference(avatar_split_clause, [], [f15986, f905, f695])).
fof(f15986, plain, (~ (e2 = op(e4, e2)) | ~ spl10_63), inference(backward_demodulation, [], [f382, f907])).
fof(f15984, plain, (~ spl10_57 | ~ spl10_67), inference(avatar_split_clause, [], [f15976, f922, f880])).
fof(f880, plain, (spl10_57 <=> (e1 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl10_57])])).
fof(f15976, plain, (~ (e1 = op(e2, e3)) | ~ spl10_67), inference(backward_demodulation, [], [f429, f924])).
fof(f924, plain, ((e1 = op(e2, e1)) | ~ spl10_67), inference(avatar_component_clause, [], [f922])).
fof(f15983, plain, (~ spl10_17 | ~ spl10_67), inference(avatar_split_clause, [], [f15974, f922, f712])).
fof(f15974, plain, (~ (e1 = op(e4, e1)) | ~ spl10_67), inference(backward_demodulation, [], [f372, f924])).
fof(f15971, plain, (~ spl10_66 | ~ spl10_71), inference(avatar_split_clause, [], [f15960, f939, f918])).
fof(f918, plain, (spl10_66 <=> (e0 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl10_66])])).
fof(f939, plain, (spl10_71 <=> (e0 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl10_71])])).
fof(f15960, plain, (~ (e0 = op(e2, e1)) | ~ spl10_71), inference(backward_demodulation, [], [f424, f941])).
fof(f941, plain, ((e0 = op(e2, e0)) | ~ spl10_71), inference(avatar_component_clause, [], [f939])).
fof(f15970, plain, (~ spl10_21 | ~ spl10_71), inference(avatar_split_clause, [], [f15959, f939, f729])).
fof(f15959, plain, (~ (e0 = op(e4, e0)) | ~ spl10_71), inference(backward_demodulation, [], [f362, f941])).
fof(f362, plain, ~ (op(e2, e0) = op(e4, e0)), inference(cnf_transformation, [], [f4])).
fof(f15949, plain, (~ spl10_35 | ~ spl10_85), inference(avatar_split_clause, [], [f15945, f997, f787])).
fof(f15945, plain, (~ (e4 = op(e3, e3)) | ~ spl10_85), inference(backward_demodulation, [], [f389, f999])).
fof(f15938, plain, (~ spl10_77 | ~ spl10_87), inference(avatar_split_clause, [], [f15929, f1006, f964])).
fof(f1006, plain, (spl10_87 <=> (e1 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl10_87])])).
fof(f15929, plain, (~ (e1 = op(e1, e4)) | ~ spl10_87), inference(backward_demodulation, [], [f422, f1008])).
fof(f1008, plain, ((e1 = op(e1, e2)) | ~ spl10_87), inference(avatar_component_clause, [], [f1006])).
fof(f15923, plain, (~ spl10_90 | ~ spl10_93 | spl10_201), inference(avatar_split_clause, [], [f15917, f1588, f1031, f1018])).
fof(f1018, plain, (spl10_90 <=> (e4 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl10_90])])).
fof(f1588, plain, (spl10_201 <=> (e4 = op(e1, op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl10_201])])).
fof(f15917, plain, (~ (e4 = op(e1, e2)) | (~ spl10_93 | spl10_201)), inference(backward_demodulation, [], [f1590, f1033])).
fof(f1590, plain, (~ (e4 = op(e1, op(e1, e1))) | spl10_201), inference(avatar_component_clause, [], [f1588])).
fof(f15922, plain, (~ spl10_43 | ~ spl10_93), inference(avatar_split_clause, [], [f15907, f1031, f821])).
fof(f821, plain, (spl10_43 <=> (e2 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl10_43])])).
fof(f15907, plain, (~ (e2 = op(e3, e1)) | ~ spl10_93), inference(backward_demodulation, [], [f369, f1033])).
fof(f369, plain, ~ (op(e1, e1) = op(e3, e1)), inference(cnf_transformation, [], [f4])).
fof(f15905, plain, (~ spl10_81 | ~ spl10_99 | spl10_156), inference(avatar_split_clause, [], [f15903, f1358, f1056, f981])).
fof(f981, plain, (spl10_81 <=> (e0 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl10_81])])).
fof(f1358, plain, (spl10_156 <=> (e0 = op(e1, op(e1, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl10_156])])).
fof(f15903, plain, (~ (e0 = op(e1, e3)) | (~ spl10_99 | spl10_156)), inference(backward_demodulation, [], [f1360, f1058])).
fof(f1360, plain, (~ (e0 = op(e1, op(e1, e0))) | spl10_156), inference(avatar_component_clause, [], [f1358])).
fof(f15892, plain, (spl10_6 | ~ spl10_104 | ~ spl10_135), inference(avatar_split_clause, [], [f15889, f1259, f1077, f666])).
fof(f1259, plain, (spl10_135 <=> (e0 = op(e4, op(e0, e4)))), introduced(avatar_definition, [new_symbols(naming, [spl10_135])])).
fof(f15889, plain, ((e0 = op(e4, e3)) | (~ spl10_104 | ~ spl10_135)), inference(backward_demodulation, [], [f1261, f1079])).
fof(f1079, plain, ((e3 = op(e0, e4)) | ~ spl10_104), inference(avatar_component_clause, [], [f1077])).
fof(f1261, plain, ((e0 = op(e4, op(e0, e4))) | ~ spl10_135), inference(avatar_component_clause, [], [f1259])).
fof(f15891, plain, (~ spl10_29 | ~ spl10_104), inference(avatar_split_clause, [], [f15886, f1077, f762])).
fof(f15886, plain, (~ (e3 = op(e3, e4)) | ~ spl10_104), inference(backward_demodulation, [], [f396, f1079])).
fof(f396, plain, ~ (op(e0, e4) = op(e3, e4)), inference(cnf_transformation, [], [f4])).
fof(f15881, plain, (~ spl10_101 | ~ spl10_111), inference(avatar_split_clause, [], [f15873, f1107, f1065])).
fof(f1107, plain, (spl10_111 <=> (e0 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl10_111])])).
fof(f15873, plain, (~ (e0 = op(e0, e4)) | ~ spl10_111), inference(backward_demodulation, [], [f412, f1109])).
fof(f1109, plain, ((e0 = op(e0, e2)) | ~ spl10_111), inference(avatar_component_clause, [], [f1107])).
fof(f15880, plain, (~ spl10_36 | ~ spl10_111), inference(avatar_split_clause, [], [f15871, f1107, f792])).
fof(f15871, plain, (~ (e0 = op(e3, e2)) | ~ spl10_111), inference(backward_demodulation, [], [f376, f1109])).
fof(f15865, plain, (~ spl10_93 | ~ spl10_122 | spl10_234), inference(avatar_split_clause, [], [f15857, f1767, f1153, f1031])).
fof(f15857, plain, (~ (e2 = op(e1, e1)) | (~ spl10_122 | spl10_234)), inference(backward_demodulation, [], [f1769, f1155])).
fof(f15864, plain, (~ spl10_112 | ~ spl10_122), inference(avatar_split_clause, [], [f15849, f1153, f1111])).
fof(f1111, plain, (spl10_112 <=> (e1 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl10_112])])).
fof(f15849, plain, (~ (e1 = op(e0, e2)) | ~ spl10_122), inference(backward_demodulation, [], [f405, f1155])).
fof(f15863, plain, (~ spl10_47 | ~ spl10_122), inference(avatar_split_clause, [], [f15847, f1153, f838])).
fof(f838, plain, (spl10_47 <=> (e1 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl10_47])])).
fof(f15847, plain, (~ (e1 = op(e3, e0)) | ~ spl10_122), inference(backward_demodulation, [], [f356, f1155])).
fof(f356, plain, ~ (op(e0, e0) = op(e3, e0)), inference(cnf_transformation, [], [f4])).
fof(f15844, plain, (spl10_15 | ~ spl10_128), inference(avatar_split_clause, [], [f15832, f1178, f703])).
fof(f1178, plain, (spl10_128 <=> (e2 = unit)), introduced(avatar_definition, [new_symbols(naming, [spl10_128])])).
fof(f15832, plain, ((e4 = op(e4, e2)) | ~ spl10_128), inference(backward_demodulation, [], [f302, f1180])).
fof(f1180, plain, ((e2 = unit) | ~ spl10_128), inference(avatar_component_clause, [], [f1178])).
fof(f15843, plain, (spl10_55 | ~ spl10_128), inference(avatar_split_clause, [], [f15831, f1178, f871])).
fof(f15831, plain, ((e4 = op(e2, e4)) | ~ spl10_128), inference(backward_demodulation, [], [f301, f1180])).
fof(f301, plain, (e4 = op(unit, e4)), inference(cnf_transformation, [], [f2])).
fof(f15842, plain, (spl10_39 | ~ spl10_128), inference(avatar_split_clause, [], [f15830, f1178, f804])).
fof(f15830, plain, ((e3 = op(e3, e2)) | ~ spl10_128), inference(backward_demodulation, [], [f300, f1180])).
fof(f15841, plain, (spl10_59 | ~ spl10_128), inference(avatar_split_clause, [], [f15829, f1178, f888])).
fof(f15829, plain, ((e3 = op(e2, e3)) | ~ spl10_128), inference(backward_demodulation, [], [f299, f1180])).
fof(f15840, plain, (spl10_63 | ~ spl10_128), inference(avatar_split_clause, [], [f15828, f1178, f905])).
fof(f15828, plain, ((e2 = op(e2, e2)) | ~ spl10_128), inference(backward_demodulation, [], [f298, f1180])).
fof(f15838, plain, (spl10_87 | ~ spl10_128), inference(avatar_split_clause, [], [f15826, f1178, f1006])).
fof(f15826, plain, ((e1 = op(e1, e2)) | ~ spl10_128), inference(backward_demodulation, [], [f296, f1180])).
fof(f15837, plain, (spl10_67 | ~ spl10_128), inference(avatar_split_clause, [], [f15825, f1178, f922])).
fof(f15825, plain, ((e1 = op(e2, e1)) | ~ spl10_128), inference(backward_demodulation, [], [f295, f1180])).
fof(f15836, plain, (spl10_111 | ~ spl10_128), inference(avatar_split_clause, [], [f15824, f1178, f1107])).
fof(f15824, plain, ((e0 = op(e0, e2)) | ~ spl10_128), inference(backward_demodulation, [], [f294, f1180])).
fof(f15835, plain, (spl10_71 | ~ spl10_128), inference(avatar_split_clause, [], [f15823, f1178, f939])).
fof(f15823, plain, ((e0 = op(e2, e0)) | ~ spl10_128), inference(backward_demodulation, [], [f293, f1180])).
fof(f15822, plain, (~ spl10_36 | ~ spl10_108 | spl10_141), inference(avatar_split_clause, [], [f15821, f1287, f1094, f792])).
fof(f1287, plain, (spl10_141 <=> (e0 = op(e3, op(e0, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl10_141])])).
fof(f15821, plain, (~ (e0 = op(e3, e2)) | (~ spl10_108 | spl10_141)), inference(forward_demodulation, [], [f1288, f1096])).
fof(f1096, plain, ((e2 = op(e0, e3)) | ~ spl10_108), inference(avatar_component_clause, [], [f1094])).
fof(f1288, plain, (~ (e0 = op(e3, op(e0, e3))) | spl10_141), inference(avatar_component_clause, [], [f1287])).
fof(f15820, plain, (spl10_76 | ~ spl10_120 | ~ spl10_152), inference(avatar_split_clause, [], [f15819, f1339, f1144, f960])).
fof(f15819, plain, ((e0 = op(e1, e4)) | (~ spl10_120 | ~ spl10_152)), inference(forward_demodulation, [], [f1341, f1146])).
fof(f1341, plain, ((e0 = op(e1, op(e0, e1))) | ~ spl10_152), inference(avatar_component_clause, [], [f1339])).
fof(f15797, plain, (~ spl10_113 | ~ spl10_108), inference(avatar_split_clause, [], [f15796, f1094, f1115])).
fof(f1115, plain, (spl10_113 <=> (e2 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl10_113])])).
fof(f15796, plain, (~ (e2 = op(e0, e2)) | ~ spl10_108), inference(forward_demodulation, [], [f411, f1096])).
fof(f411, plain, ~ (op(e0, e2) = op(e0, e3)), inference(cnf_transformation, [], [f4])).
fof(f15795, plain, (~ spl10_115 | ~ spl10_120), inference(avatar_split_clause, [], [f15794, f1144, f1123])).
fof(f15794, plain, (~ (e4 = op(e0, e2)) | ~ spl10_120), inference(forward_demodulation, [], [f408, f1146])).
fof(f408, plain, ~ (op(e0, e1) = op(e0, e2)), inference(cnf_transformation, [], [f4])).
fof(f15793, plain, (~ spl10_114 | ~ spl10_108 | spl10_142), inference(avatar_split_clause, [], [f15700, f1292, f1094, f1119])).
fof(f1292, plain, (spl10_142 <=> (e3 = op(e0, op(e0, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl10_142])])).
fof(f15700, plain, (~ (e3 = op(e0, e2)) | (~ spl10_108 | spl10_142)), inference(backward_demodulation, [], [f1294, f1096])).
fof(f1294, plain, (~ (e3 = op(e0, op(e0, e3))) | spl10_142), inference(avatar_component_clause, [], [f1292])).
fof(f15792, plain, (~ spl10_103 | ~ spl10_108), inference(avatar_split_clause, [], [f15791, f1094, f1073])).
fof(f1073, plain, (spl10_103 <=> (e2 = op(e0, e4))), introduced(avatar_definition, [new_symbols(naming, [spl10_103])])).
fof(f15791, plain, (~ (e2 = op(e0, e4)) | ~ spl10_108), inference(forward_demodulation, [], [f413, f1096])).
fof(f413, plain, ~ (op(e0, e3) = op(e0, e4)), inference(cnf_transformation, [], [f4])).
fof(f15790, plain, (~ spl10_105 | ~ spl10_120), inference(avatar_split_clause, [], [f15789, f1144, f1081])).
fof(f1081, plain, (spl10_105 <=> (e4 = op(e0, e4))), introduced(avatar_definition, [new_symbols(naming, [spl10_105])])).
fof(f15789, plain, (~ (e4 = op(e0, e4)) | ~ spl10_120), inference(forward_demodulation, [], [f410, f1146])).
fof(f410, plain, ~ (op(e0, e1) = op(e0, e4)), inference(cnf_transformation, [], [f4])).
fof(f15786, plain, (~ spl10_70 | ~ spl10_120), inference(avatar_split_clause, [], [f15785, f1144, f934])).
fof(f15785, plain, (~ (e4 = op(e2, e1)) | ~ spl10_120), inference(forward_demodulation, [], [f365, f1146])).
fof(f365, plain, ~ (op(e0, e1) = op(e2, e1)), inference(cnf_transformation, [], [f4])).
fof(f15784, plain, (~ spl10_58 | ~ spl10_108), inference(avatar_split_clause, [], [f15783, f1094, f884])).
fof(f884, plain, (spl10_58 <=> (e2 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl10_58])])).
fof(f15783, plain, (~ (e2 = op(e2, e3)) | ~ spl10_108), inference(forward_demodulation, [], [f385, f1096])).
fof(f385, plain, ~ (op(e0, e3) = op(e2, e3)), inference(cnf_transformation, [], [f4])).
fof(f15782, plain, (~ spl10_20 | ~ spl10_120), inference(avatar_split_clause, [], [f15781, f1144, f724])).
fof(f724, plain, (spl10_20 <=> (e4 = op(e4, e1))), introduced(avatar_definition, [new_symbols(naming, [spl10_20])])).
fof(f15781, plain, (~ (e4 = op(e4, e1)) | ~ spl10_120), inference(forward_demodulation, [], [f367, f1146])).
fof(f367, plain, ~ (op(e0, e1) = op(e4, e1)), inference(cnf_transformation, [], [f4])).
fof(f15773, plain, (~ spl10_1 | ~ spl10_35 | spl10_190), inference(avatar_split_clause, [], [f15761, f1539, f787, f645])).
fof(f645, plain, (spl10_1 <=> (e0 = op(e4, e4))), introduced(avatar_definition, [new_symbols(naming, [spl10_1])])).
fof(f1539, plain, (spl10_190 <=> (e0 = op(op(e3, e3), op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl10_190])])).
fof(f15761, plain, (~ (e0 = op(e4, e4)) | (~ spl10_35 | spl10_190)), inference(backward_demodulation, [], [f1541, f789])).
fof(f789, plain, ((e4 = op(e3, e3)) | ~ spl10_35), inference(avatar_component_clause, [], [f787])).
fof(f1541, plain, (~ (e0 = op(op(e3, e3), op(e3, e3))) | spl10_190), inference(avatar_component_clause, [], [f1539])).
fof(f15772, plain, (~ spl10_29 | ~ spl10_35 | spl10_165), inference(avatar_contradiction_clause, [], [f15771])).
fof(f15771, plain, ($false | (~ spl10_29 | ~ spl10_35 | spl10_165)), inference(subsumption_resolution, [], [f15760, f764])).
fof(f15760, plain, (~ (e3 = op(e3, e4)) | (~ spl10_35 | spl10_165)), inference(backward_demodulation, [], [f1406, f789])).
fof(f1406, plain, (~ (e3 = op(e3, op(e3, e3))) | spl10_165), inference(avatar_component_clause, [], [f1405])).
fof(f1405, plain, (spl10_165 <=> (e3 = op(e3, op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl10_165])])).
fof(f15759, plain, (~ spl10_33 | ~ spl10_43), inference(avatar_split_clause, [], [f15755, f821, f779])).
fof(f15755, plain, (~ (e2 = op(e3, e3)) | ~ spl10_43), inference(backward_demodulation, [], [f439, f823])).
fof(f823, plain, ((e2 = op(e3, e1)) | ~ spl10_43), inference(avatar_component_clause, [], [f821])).
fof(f439, plain, ~ (op(e3, e1) = op(e3, e3)), inference(cnf_transformation, [], [f4])).
fof(f15752, plain, (~ spl10_64 | ~ spl10_81 | ~ spl10_209 | spl10_215), inference(avatar_contradiction_clause, [], [f15751])).
fof(f15751, plain, ($false | (~ spl10_64 | ~ spl10_81 | ~ spl10_209 | spl10_215)), inference(subsumption_resolution, [], [f15747, f983])).
fof(f983, plain, ((e0 = op(e1, e3)) | ~ spl10_81), inference(avatar_component_clause, [], [f981])).
fof(f15747, plain, (~ (e0 = op(e1, e3)) | (~ spl10_64 | ~ spl10_209 | spl10_215)), inference(backward_demodulation, [], [f15642, f911])).
fof(f15642, plain, (~ (e0 = op(e1, op(e2, e2))) | (~ spl10_209 | spl10_215)), inference(backward_demodulation, [], [f1663, f1628])).
fof(f1628, plain, ((e1 = op(e2, op(e2, e2))) | ~ spl10_209), inference(avatar_component_clause, [], [f1627])).
fof(f1627, plain, (spl10_209 <=> (e1 = op(e2, op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl10_209])])).
fof(f1663, plain, (~ (e0 = op(op(e2, op(e2, e2)), op(e2, e2))) | spl10_215), inference(avatar_component_clause, [], [f1661])).
fof(f1661, plain, (spl10_215 <=> (e0 = op(op(e2, op(e2, e2)), op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl10_215])])).
fof(f15738, plain, (~ spl10_65 | ~ spl10_75), inference(avatar_split_clause, [], [f15735, f955, f913])).
fof(f15735, plain, (~ (e4 = op(e2, e2)) | ~ spl10_75), inference(backward_demodulation, [], [f425, f957])).
fof(f425, plain, ~ (op(e2, e0) = op(e2, e2)), inference(cnf_transformation, [], [f4])).
fof(f15733, plain, (~ spl10_65 | ~ spl10_90), inference(avatar_split_clause, [], [f15730, f1018, f913])).
fof(f15730, plain, (~ (e4 = op(e2, e2)) | ~ spl10_90), inference(backward_demodulation, [], [f378, f1020])).
fof(f1020, plain, ((e4 = op(e1, e2)) | ~ spl10_90), inference(avatar_component_clause, [], [f1018])).
fof(f378, plain, ~ (op(e1, e2) = op(e2, e2)), inference(cnf_transformation, [], [f4])).
fof(f15729, plain, (~ spl10_94 | ~ spl10_108 | spl10_199 | ~ spl10_237), inference(avatar_contradiction_clause, [], [f15728])).
fof(f15728, plain, ($false | (~ spl10_94 | ~ spl10_108 | spl10_199 | ~ spl10_237)), inference(subsumption_resolution, [], [f15721, f1096])).
fof(f15721, plain, (~ (e2 = op(e0, e3)) | (~ spl10_94 | spl10_199 | ~ spl10_237)), inference(backward_demodulation, [], [f15647, f1037])).
fof(f15647, plain, (~ (e2 = op(e0, op(e1, e1))) | (spl10_199 | ~ spl10_237)), inference(forward_demodulation, [], [f1581, f1788])).
fof(f1788, plain, ((e0 = op(e1, op(e1, e1))) | ~ spl10_237), inference(avatar_component_clause, [], [f1787])).
fof(f1581, plain, (~ (e2 = op(op(e1, op(e1, e1)), op(e1, e1))) | spl10_199), inference(avatar_component_clause, [], [f1579])).
fof(f1579, plain, (spl10_199 <=> (e2 = op(op(e1, op(e1, e1)), op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl10_199])])).
fof(f15722, plain, (~ spl10_89 | ~ spl10_94), inference(avatar_split_clause, [], [f15712, f1035, f1014])).
fof(f1014, plain, (spl10_89 <=> (e3 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl10_89])])).
fof(f15712, plain, (~ (e3 = op(e1, e2)) | ~ spl10_94), inference(backward_demodulation, [], [f418, f1037])).
fof(f418, plain, ~ (op(e1, e1) = op(e1, e2)), inference(cnf_transformation, [], [f4])).
fof(f15710, plain, (~ spl10_93 | ~ spl10_98), inference(avatar_split_clause, [], [f15704, f1052, f1031])).
fof(f1052, plain, (spl10_98 <=> (e2 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl10_98])])).
fof(f15704, plain, (~ (e2 = op(e1, e1)) | ~ spl10_98), inference(backward_demodulation, [], [f414, f1054])).
fof(f1054, plain, ((e2 = op(e1, e0)) | ~ spl10_98), inference(avatar_component_clause, [], [f1052])).
fof(f15701, plain, (~ spl10_33 | ~ spl10_108), inference(avatar_split_clause, [], [f15696, f1094, f779])).
fof(f15696, plain, (~ (e2 = op(e3, e3)) | ~ spl10_108), inference(backward_demodulation, [], [f386, f1096])).
fof(f15694, plain, (~ spl10_110 | ~ spl10_120), inference(avatar_split_clause, [], [f15690, f1144, f1102])).
fof(f15690, plain, (~ (e4 = op(e0, e3)) | ~ spl10_120), inference(backward_demodulation, [], [f409, f1146])).
fof(f15693, plain, (~ spl10_45 | ~ spl10_120), inference(avatar_split_clause, [], [f15689, f1144, f829])).
fof(f829, plain, (spl10_45 <=> (e4 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl10_45])])).
fof(f15689, plain, (~ (e4 = op(e3, e1)) | ~ spl10_120), inference(backward_demodulation, [], [f366, f1146])).
fof(f366, plain, ~ (op(e0, e1) = op(e3, e1)), inference(cnf_transformation, [], [f4])).
fof(f15687, plain, (~ spl10_57 | ~ spl10_124 | ~ spl10_227 | spl10_235), inference(avatar_contradiction_clause, [], [f15686])).
fof(f15686, plain, ($false | (~ spl10_57 | ~ spl10_124 | ~ spl10_227 | spl10_235)), inference(subsumption_resolution, [], [f15680, f882])).
fof(f882, plain, ((e1 = op(e2, e3)) | ~ spl10_57), inference(avatar_component_clause, [], [f880])).
fof(f15680, plain, (~ (e1 = op(e2, e3)) | (~ spl10_124 | ~ spl10_227 | spl10_235)), inference(backward_demodulation, [], [f15637, f1163])).
fof(f15637, plain, (~ (e1 = op(e2, op(e0, e0))) | (~ spl10_227 | spl10_235)), inference(backward_demodulation, [], [f1773, f1722])).
fof(f15683, plain, (spl10_108 | ~ spl10_124 | ~ spl10_227), inference(avatar_split_clause, [], [f15675, f1721, f1161, f1094])).
fof(f15675, plain, ((e2 = op(e0, e3)) | (~ spl10_124 | ~ spl10_227)), inference(backward_demodulation, [], [f1722, f1163])).
fof(f15682, plain, (~ spl10_119 | ~ spl10_124), inference(avatar_split_clause, [], [f15671, f1161, f1140])).
fof(f1140, plain, (spl10_119 <=> (e3 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl10_119])])).
fof(f15671, plain, (~ (e3 = op(e0, e1)) | ~ spl10_124), inference(backward_demodulation, [], [f404, f1163])).
fof(f404, plain, ~ (op(e0, e0) = op(e0, e1)), inference(cnf_transformation, [], [f4])).
fof(f15681, plain, (~ spl10_74 | ~ spl10_124), inference(avatar_split_clause, [], [f15670, f1161, f951])).
fof(f951, plain, (spl10_74 <=> (e3 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl10_74])])).
fof(f15670, plain, (~ (e3 = op(e2, e0)) | ~ spl10_124), inference(backward_demodulation, [], [f355, f1163])).
fof(f15663, plain, (spl10_43 | ~ spl10_57 | ~ spl10_164), inference(avatar_split_clause, [], [f15662, f1401, f880, f821])).
fof(f15662, plain, ((e2 = op(e3, e1)) | (~ spl10_57 | ~ spl10_164)), inference(forward_demodulation, [], [f1403, f882])).
fof(f15656, plain, (~ spl10_74 | ~ spl10_36 | spl10_170), inference(avatar_split_clause, [], [f15655, f1433, f792, f951])).
fof(f1433, plain, (spl10_170 <=> (e3 = op(e2, op(e3, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl10_170])])).
fof(f15655, plain, (~ (e3 = op(e2, e0)) | (~ spl10_36 | spl10_170)), inference(forward_demodulation, [], [f1434, f794])).
fof(f794, plain, ((e0 = op(e3, e2)) | ~ spl10_36), inference(avatar_component_clause, [], [f792])).
fof(f1434, plain, (~ (e3 = op(e2, op(e3, e2))) | spl10_170), inference(avatar_component_clause, [], [f1433])).
fof(f15652, plain, (spl10_98 | ~ spl10_66 | ~ spl10_174), inference(avatar_split_clause, [], [f15651, f1457, f918, f1052])).
fof(f1457, plain, (spl10_174 <=> (e2 = op(e1, op(e2, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl10_174])])).
fof(f15651, plain, ((e2 = op(e1, e0)) | (~ spl10_66 | ~ spl10_174)), inference(forward_demodulation, [], [f1459, f920])).
fof(f920, plain, ((e0 = op(e2, e1)) | ~ spl10_66), inference(avatar_component_clause, [], [f918])).
fof(f1459, plain, ((e2 = op(e1, op(e2, e1))) | ~ spl10_174), inference(avatar_component_clause, [], [f1457])).
fof(f15581, plain, (~ spl10_1 | ~ spl10_5), inference(avatar_contradiction_clause, [], [f15580])).
fof(f15580, plain, ($false | (~ spl10_1 | ~ spl10_5)), inference(subsumption_resolution, [], [f15579, f457])).
fof(f457, plain, ~ (e0 = e4), inference(cnf_transformation, [], [f5])).
fof(f5, plain, (~ (e3 = e4) & ~ (e2 = e4) & ~ (e2 = e3) & ~ (e1 = e4) & ~ (e1 = e3) & ~ (e1 = e2) & ~ (e0 = e4) & ~ (e0 = e3) & ~ (e0 = e2) & ~ (e0 = e1)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax5)).
fof(f15579, plain, ((e0 = e4) | (~ spl10_1 | ~ spl10_5)), inference(backward_demodulation, [], [f663, f647])).
fof(f647, plain, ((e0 = op(e4, e4)) | ~ spl10_1), inference(avatar_component_clause, [], [f645])).
fof(f663, plain, ((e4 = op(e4, e4)) | ~ spl10_5), inference(avatar_component_clause, [], [f661])).
fof(f15557, plain, (~ spl10_8 | ~ spl10_13), inference(avatar_split_clause, [], [f15552, f695, f674])).
fof(f674, plain, (spl10_8 <=> (e2 = op(e4, e3))), introduced(avatar_definition, [new_symbols(naming, [spl10_8])])).
fof(f15552, plain, (~ (e2 = op(e4, e3)) | ~ spl10_13), inference(backward_demodulation, [], [f451, f697])).
fof(f697, plain, ((e2 = op(e4, e2)) | ~ spl10_13), inference(avatar_component_clause, [], [f695])).
fof(f451, plain, ~ (op(e4, e2) = op(e4, e3)), inference(cnf_transformation, [], [f4])).
fof(f15522, plain, (~ spl10_8 | ~ spl10_33), inference(avatar_split_clause, [], [f15509, f779, f674])).
fof(f15509, plain, (~ (e2 = op(e4, e3)) | ~ spl10_33), inference(backward_demodulation, [], [f393, f781])).
fof(f393, plain, ~ (op(e3, e3) = op(e4, e3)), inference(cnf_transformation, [], [f4])).
fof(f15508, plain, (~ spl10_26 | ~ spl10_36), inference(avatar_split_clause, [], [f15503, f792, f750])).
fof(f750, plain, (spl10_26 <=> (e0 = op(e3, e4))), introduced(avatar_definition, [new_symbols(naming, [spl10_26])])).
fof(f15503, plain, (~ (e0 = op(e3, e4)) | ~ spl10_36), inference(backward_demodulation, [], [f442, f794])).
fof(f15500, plain, (~ spl10_32 | ~ spl10_47), inference(avatar_split_clause, [], [f15493, f838, f775])).
fof(f15493, plain, (~ (e1 = op(e3, e3)) | ~ spl10_47), inference(backward_demodulation, [], [f436, f840])).
fof(f840, plain, ((e1 = op(e3, e0)) | ~ spl10_47), inference(avatar_component_clause, [], [f838])).
fof(f15475, plain, (~ spl10_32 | ~ spl10_57), inference(avatar_split_clause, [], [f15467, f880, f775])).
fof(f15467, plain, (~ (e1 = op(e3, e3)) | ~ spl10_57), inference(backward_demodulation, [], [f391, f882])).
fof(f391, plain, ~ (op(e2, e3) = op(e3, e3)), inference(cnf_transformation, [], [f4])).
fof(f15466, plain, (spl10_1 | ~ spl10_65 | ~ spl10_186), inference(avatar_split_clause, [], [f15457, f1521, f913, f645])).
fof(f1521, plain, (spl10_186 <=> (e0 = op(op(e2, e2), op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl10_186])])).
fof(f15457, plain, ((e0 = op(e4, e4)) | (~ spl10_65 | ~ spl10_186)), inference(backward_demodulation, [], [f1522, f915])).
fof(f1522, plain, ((e0 = op(op(e2, e2), op(e2, e2))) | ~ spl10_186), inference(avatar_component_clause, [], [f1521])).
fof(f15452, plain, (~ spl10_16 | ~ spl10_66), inference(avatar_split_clause, [], [f15442, f918, f708])).
fof(f708, plain, (spl10_16 <=> (e0 = op(e4, e1))), introduced(avatar_definition, [new_symbols(naming, [spl10_16])])).
fof(f15442, plain, (~ (e0 = op(e4, e1)) | ~ spl10_66), inference(backward_demodulation, [], [f372, f920])).
fof(f15441, plain, (~ spl10_59 | ~ spl10_74), inference(avatar_split_clause, [], [f15435, f951, f888])).
fof(f15435, plain, (~ (e3 = op(e2, e3)) | ~ spl10_74), inference(backward_demodulation, [], [f426, f953])).
fof(f953, plain, ((e3 = op(e2, e0)) | ~ spl10_74), inference(avatar_component_clause, [], [f951])).
fof(f426, plain, ~ (op(e2, e0) = op(e2, e3)), inference(cnf_transformation, [], [f4])).
fof(f15440, plain, (~ spl10_24 | ~ spl10_74), inference(avatar_split_clause, [], [f15432, f951, f741])).
fof(f741, plain, (spl10_24 <=> (e3 = op(e4, e0))), introduced(avatar_definition, [new_symbols(naming, [spl10_24])])).
fof(f15432, plain, (~ (e3 = op(e4, e0)) | ~ spl10_74), inference(backward_demodulation, [], [f362, f953])).
fof(f15420, plain, (spl10_57 | ~ spl10_89 | ~ spl10_144), inference(avatar_split_clause, [], [f15417, f1301, f1014, f880])).
fof(f15417, plain, ((e1 = op(e2, e3)) | (~ spl10_89 | ~ spl10_144)), inference(backward_demodulation, [], [f1303, f1016])).
fof(f1016, plain, ((e3 = op(e1, e2)) | ~ spl10_89), inference(avatar_component_clause, [], [f1014])).
fof(f15419, plain, (~ spl10_39 | ~ spl10_89), inference(avatar_split_clause, [], [f15413, f1014, f804])).
fof(f15413, plain, (~ (e3 = op(e3, e2)) | ~ spl10_89), inference(backward_demodulation, [], [f379, f1016])).
fof(f379, plain, ~ (op(e1, e2) = op(e3, e2)), inference(cnf_transformation, [], [f4])).
fof(f15409, plain, (~ spl10_36 | ~ spl10_93 | ~ spl10_200 | spl10_231), inference(avatar_split_clause, [], [f15403, f1749, f1583, f1031, f792])).
fof(f1583, plain, (spl10_200 <=> (e3 = op(e1, op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl10_200])])).
fof(f15403, plain, (~ (e0 = op(e3, e2)) | (~ spl10_93 | ~ spl10_200 | spl10_231)), inference(backward_demodulation, [], [f14793, f1033])).
fof(f14793, plain, (~ (e0 = op(e3, op(e1, e1))) | (~ spl10_200 | spl10_231)), inference(forward_demodulation, [], [f1751, f1584])).
fof(f1584, plain, ((e3 = op(e1, op(e1, e1))) | ~ spl10_200), inference(avatar_component_clause, [], [f1583])).
fof(f15406, plain, (spl10_89 | ~ spl10_93 | ~ spl10_200), inference(avatar_split_clause, [], [f15397, f1583, f1031, f1014])).
fof(f15397, plain, ((e3 = op(e1, e2)) | (~ spl10_93 | ~ spl10_200)), inference(backward_demodulation, [], [f1584, f1033])).
fof(f15405, plain, (~ spl10_78 | ~ spl10_93), inference(avatar_split_clause, [], [f15393, f1031, f968])).
fof(f968, plain, (spl10_78 <=> (e2 = op(e1, e4))), introduced(avatar_definition, [new_symbols(naming, [spl10_78])])).
fof(f15393, plain, (~ (e2 = op(e1, e4)) | ~ spl10_93), inference(backward_demodulation, [], [f420, f1033])).
fof(f420, plain, ~ (op(e1, e1) = op(e1, e4)), inference(cnf_transformation, [], [f4])).
fof(f15386, plain, (~ spl10_26 | ~ spl10_101), inference(avatar_split_clause, [], [f15378, f1065, f750])).
fof(f15378, plain, (~ (e0 = op(e3, e4)) | ~ spl10_101), inference(backward_demodulation, [], [f396, f1067])).
fof(f15375, plain, (spl10_66 | ~ spl10_112 | ~ spl10_147), inference(avatar_split_clause, [], [f15372, f1315, f1111, f918])).
fof(f15372, plain, ((e0 = op(e2, e1)) | (~ spl10_112 | ~ spl10_147)), inference(backward_demodulation, [], [f1317, f1113])).
fof(f1113, plain, ((e1 = op(e0, e2)) | ~ spl10_112), inference(avatar_component_clause, [], [f1111])).
fof(f15374, plain, (~ spl10_87 | ~ spl10_112), inference(avatar_split_clause, [], [f15364, f1111, f1006])).
fof(f15364, plain, (~ (e1 = op(e1, e2)) | ~ spl10_112), inference(backward_demodulation, [], [f374, f1113])).
fof(f374, plain, ~ (op(e0, e2) = op(e1, e2)), inference(cnf_transformation, [], [f4])).
fof(f15363, plain, (~ spl10_104 | ~ spl10_119), inference(avatar_split_clause, [], [f15358, f1140, f1077])).
fof(f15358, plain, (~ (e3 = op(e0, e4)) | ~ spl10_119), inference(backward_demodulation, [], [f410, f1142])).
fof(f1142, plain, ((e3 = op(e0, e1)) | ~ spl10_119), inference(avatar_component_clause, [], [f1140])).
fof(f15362, plain, (~ spl10_94 | ~ spl10_119), inference(avatar_split_clause, [], [f15354, f1140, f1035])).
fof(f15354, plain, (~ (e3 = op(e1, e1)) | ~ spl10_119), inference(backward_demodulation, [], [f364, f1142])).
fof(f15353, plain, (~ spl10_89 | ~ spl10_123 | spl10_226 | ~ spl10_236), inference(avatar_split_clause, [], [f15347, f1781, f1717, f1157, f1014])).
fof(f15347, plain, (~ (e3 = op(e1, e2)) | (~ spl10_123 | spl10_226 | ~ spl10_236)), inference(backward_demodulation, [], [f15294, f1159])).
fof(f15294, plain, (~ (e3 = op(e1, op(e0, e0))) | (spl10_226 | ~ spl10_236)), inference(backward_demodulation, [], [f1719, f1782])).
fof(f15351, plain, (spl10_112 | ~ spl10_123 | ~ spl10_236), inference(avatar_split_clause, [], [f15344, f1781, f1157, f1111])).
fof(f15344, plain, ((e1 = op(e0, e2)) | (~ spl10_123 | ~ spl10_236)), inference(backward_demodulation, [], [f1782, f1159])).
fof(f15349, plain, (~ spl10_118 | ~ spl10_123), inference(avatar_split_clause, [], [f15337, f1157, f1136])).
fof(f15337, plain, (~ (e2 = op(e0, e1)) | ~ spl10_123), inference(backward_demodulation, [], [f404, f1159])).
fof(f15348, plain, (~ spl10_48 | ~ spl10_123), inference(avatar_split_clause, [], [f15335, f1157, f842])).
fof(f15335, plain, (~ (e2 = op(e3, e0)) | ~ spl10_123), inference(backward_demodulation, [], [f356, f1159])).
fof(f15313, plain, (spl10_47 | ~ spl10_81 | ~ spl10_138), inference(avatar_split_clause, [], [f15312, f1273, f981, f838])).
fof(f1273, plain, (spl10_138 <=> (e1 = op(e3, op(e1, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl10_138])])).
fof(f15312, plain, ((e1 = op(e3, e0)) | (~ spl10_81 | ~ spl10_138)), inference(forward_demodulation, [], [f1275, f983])).
fof(f1275, plain, ((e1 = op(e3, op(e1, e3))) | ~ spl10_138), inference(avatar_component_clause, [], [f1273])).
fof(f15280, plain, (~ spl10_90 | ~ spl10_100), inference(avatar_split_clause, [], [f15279, f1060, f1018])).
fof(f15279, plain, (~ (e4 = op(e1, e2)) | ~ spl10_100), inference(forward_demodulation, [], [f415, f1062])).
fof(f415, plain, ~ (op(e1, e0) = op(e1, e2)), inference(cnf_transformation, [], [f4])).
fof(f15277, plain, (~ spl10_75 | ~ spl10_100), inference(avatar_split_clause, [], [f15276, f1060, f955])).
fof(f15276, plain, (~ (e4 = op(e2, e0)) | ~ spl10_100), inference(forward_demodulation, [], [f358, f1062])).
fof(f358, plain, ~ (op(e1, e0) = op(e2, e0)), inference(cnf_transformation, [], [f4])).
fof(f15274, plain, (~ spl10_60 | ~ spl10_110), inference(avatar_split_clause, [], [f15273, f1102, f892])).
fof(f892, plain, (spl10_60 <=> (e4 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl10_60])])).
fof(f15273, plain, (~ (e4 = op(e2, e3)) | ~ spl10_110), inference(forward_demodulation, [], [f385, f1104])).
fof(f15271, plain, (~ spl10_40 | ~ spl10_45), inference(avatar_split_clause, [], [f15270, f829, f808])).
fof(f15270, plain, (~ (e4 = op(e3, e2)) | ~ spl10_45), inference(forward_demodulation, [], [f438, f831])).
fof(f831, plain, ((e4 = op(e3, e1)) | ~ spl10_45), inference(avatar_component_clause, [], [f829])).
fof(f15269, plain, (~ spl10_30 | ~ spl10_45), inference(avatar_split_clause, [], [f15268, f829, f766])).
fof(f766, plain, (spl10_30 <=> (e4 = op(e3, e4))), introduced(avatar_definition, [new_symbols(naming, [spl10_30])])).
fof(f15268, plain, (~ (e4 = op(e3, e4)) | ~ spl10_45), inference(forward_demodulation, [], [f440, f831])).
fof(f440, plain, ~ (op(e3, e1) = op(e3, e4)), inference(cnf_transformation, [], [f4])).
fof(f15267, plain, (~ spl10_25 | ~ spl10_100), inference(avatar_split_clause, [], [f15266, f1060, f745])).
fof(f745, plain, (spl10_25 <=> (e4 = op(e4, e0))), introduced(avatar_definition, [new_symbols(naming, [spl10_25])])).
fof(f15266, plain, (~ (e4 = op(e4, e0)) | ~ spl10_100), inference(forward_demodulation, [], [f360, f1062])).
fof(f360, plain, ~ (op(e1, e0) = op(e4, e0)), inference(cnf_transformation, [], [f4])).
fof(f15265, plain, (~ spl10_20 | ~ spl10_45), inference(avatar_split_clause, [], [f15264, f829, f724])).
fof(f15264, plain, (~ (e4 = op(e4, e1)) | ~ spl10_45), inference(forward_demodulation, [], [f373, f831])).
fof(f373, plain, ~ (op(e3, e1) = op(e4, e1)), inference(cnf_transformation, [], [f4])).
fof(f15263, plain, (~ spl10_2 | ~ spl10_118 | spl10_204 | ~ spl10_225), inference(avatar_contradiction_clause, [], [f15262])).
fof(f15262, plain, ($false | (~ spl10_2 | ~ spl10_118 | spl10_204 | ~ spl10_225)), inference(subsumption_resolution, [], [f15261, f1138])).
fof(f15261, plain, (~ (e2 = op(e0, e1)) | (~ spl10_2 | spl10_204 | ~ spl10_225)), inference(backward_demodulation, [], [f15136, f651])).
fof(f15136, plain, (~ (e2 = op(e0, op(e4, e4))) | (spl10_204 | ~ spl10_225)), inference(forward_demodulation, [], [f1605, f1712])).
fof(f1712, plain, ((e0 = op(e4, op(e4, e4))) | ~ spl10_225), inference(avatar_component_clause, [], [f1711])).
fof(f1711, plain, (spl10_225 <=> (e0 = op(e4, op(e4, e4)))), introduced(avatar_definition, [new_symbols(naming, [spl10_225])])).
fof(f1605, plain, (~ (e2 = op(op(e4, op(e4, e4)), op(e4, e4))) | spl10_204), inference(avatar_component_clause, [], [f1603])).
fof(f1603, plain, (spl10_204 <=> (e2 = op(op(e4, op(e4, e4)), op(e4, e4)))), introduced(avatar_definition, [new_symbols(naming, [spl10_204])])).
fof(f15251, plain, (~ spl10_3 | ~ spl10_8), inference(avatar_split_clause, [], [f15246, f674, f653])).
fof(f653, plain, (spl10_3 <=> (e2 = op(e4, e4))), introduced(avatar_definition, [new_symbols(naming, [spl10_3])])).
fof(f15246, plain, (~ (e2 = op(e4, e4)) | ~ spl10_8), inference(backward_demodulation, [], [f453, f676])).
fof(f676, plain, ((e2 = op(e4, e3)) | ~ spl10_8), inference(avatar_component_clause, [], [f674])).
fof(f453, plain, ~ (op(e4, e3) = op(e4, e4)), inference(cnf_transformation, [], [f4])).
fof(f15244, plain, (~ spl10_7 | ~ spl10_32), inference(avatar_split_clause, [], [f15231, f775, f670])).
fof(f15231, plain, (~ (e1 = op(e4, e3)) | ~ spl10_32), inference(backward_demodulation, [], [f393, f777])).
fof(f15229, plain, (~ spl10_35 | ~ spl10_45), inference(avatar_split_clause, [], [f15227, f829, f787])).
fof(f15227, plain, (~ (e4 = op(e3, e3)) | ~ spl10_45), inference(backward_demodulation, [], [f439, f831])).
fof(f15225, plain, (~ spl10_43 | ~ spl10_48), inference(avatar_split_clause, [], [f15219, f842, f821])).
fof(f15219, plain, (~ (e2 = op(e3, e1)) | ~ spl10_48), inference(backward_demodulation, [], [f434, f844])).
fof(f434, plain, ~ (op(e3, e0) = op(e3, e1)), inference(cnf_transformation, [], [f4])).
fof(f15217, plain, (~ spl10_3 | ~ spl10_78), inference(avatar_split_clause, [], [f15212, f968, f653])).
fof(f15212, plain, (~ (e2 = op(e4, e4)) | ~ spl10_78), inference(backward_demodulation, [], [f400, f970])).
fof(f970, plain, ((e2 = op(e1, e4)) | ~ spl10_78), inference(avatar_component_clause, [], [f968])).
fof(f400, plain, ~ (op(e1, e4) = op(e4, e4)), inference(cnf_transformation, [], [f4])).
fof(f15211, plain, (~ spl10_35 | ~ spl10_94 | spl10_242), inference(avatar_split_clause, [], [f15207, f1841, f1035, f787])).
fof(f15207, plain, (~ (e4 = op(e3, e3)) | (~ spl10_94 | spl10_242)), inference(backward_demodulation, [], [f1843, f1037])).
fof(f1843, plain, (~ (e4 = op(op(e1, e1), op(e1, e1))) | spl10_242), inference(avatar_component_clause, [], [f1841])).
fof(f15210, plain, (~ spl10_79 | ~ spl10_94), inference(avatar_split_clause, [], [f15201, f1035, f972])).
fof(f972, plain, (spl10_79 <=> (e3 = op(e1, e4))), introduced(avatar_definition, [new_symbols(naming, [spl10_79])])).
fof(f15201, plain, (~ (e3 = op(e1, e4)) | ~ spl10_94), inference(backward_demodulation, [], [f420, f1037])).
fof(f15199, plain, (~ spl10_95 | ~ spl10_100), inference(avatar_split_clause, [], [f15195, f1060, f1039])).
fof(f15195, plain, (~ (e4 = op(e1, e1)) | ~ spl10_100), inference(backward_demodulation, [], [f414, f1062])).
fof(f15193, plain, (~ spl10_79 | ~ spl10_104), inference(avatar_split_clause, [], [f15188, f1077, f972])).
fof(f15188, plain, (~ (e3 = op(e1, e4)) | ~ spl10_104), inference(backward_demodulation, [], [f394, f1079])).
fof(f15187, plain, (~ spl10_104 | ~ spl10_110 | spl10_142), inference(avatar_split_clause, [], [f15185, f1292, f1102, f1077])).
fof(f15185, plain, (~ (e3 = op(e0, e4)) | (~ spl10_110 | spl10_142)), inference(backward_demodulation, [], [f1294, f1104])).
fof(f15177, plain, (~ spl10_43 | ~ spl10_118), inference(avatar_split_clause, [], [f15169, f1136, f821])).
fof(f15169, plain, (~ (e2 = op(e3, e1)) | ~ spl10_118), inference(backward_demodulation, [], [f366, f1138])).
fof(f15165, plain, (~ spl10_118 | ~ spl10_122 | spl10_227), inference(avatar_split_clause, [], [f15160, f1721, f1153, f1136])).
fof(f15160, plain, (~ (e2 = op(e0, e1)) | (~ spl10_122 | spl10_227)), inference(backward_demodulation, [], [f1723, f1155])).
fof(f1723, plain, (~ (e2 = op(e0, op(e0, e0))) | spl10_227), inference(avatar_component_clause, [], [f1721])).
fof(f15164, plain, (~ spl10_119 | ~ spl10_122 | spl10_222), inference(avatar_split_clause, [], [f15157, f1695, f1153, f1140])).
fof(f15157, plain, (~ (e3 = op(e0, e1)) | (~ spl10_122 | spl10_222)), inference(backward_demodulation, [], [f1697, f1155])).
fof(f1697, plain, (~ (e3 = op(e0, op(e0, e0))) | spl10_222), inference(avatar_component_clause, [], [f1695])).
fof(f15163, plain, (~ spl10_102 | ~ spl10_122), inference(avatar_split_clause, [], [f15152, f1153, f1069])).
fof(f1069, plain, (spl10_102 <=> (e1 = op(e0, e4))), introduced(avatar_definition, [new_symbols(naming, [spl10_102])])).
fof(f15152, plain, (~ (e1 = op(e0, e4)) | ~ spl10_122), inference(backward_demodulation, [], [f407, f1155])).
fof(f407, plain, ~ (op(e0, e0) = op(e0, e4)), inference(cnf_transformation, [], [f4])).
fof(f15147, plain, (~ spl10_47 | ~ spl10_81 | spl10_138), inference(avatar_split_clause, [], [f15146, f1273, f981, f838])).
fof(f15146, plain, (~ (e1 = op(e3, e0)) | (~ spl10_81 | spl10_138)), inference(forward_demodulation, [], [f1274, f983])).
fof(f1274, plain, (~ (e1 = op(e3, op(e1, e3))) | spl10_138), inference(avatar_component_clause, [], [f1273])).
fof(f15143, plain, (spl10_100 | ~ spl10_16 | ~ spl10_176), inference(avatar_split_clause, [], [f15142, f1465, f708, f1060])).
fof(f1465, plain, (spl10_176 <=> (e4 = op(e1, op(e4, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl10_176])])).
fof(f15142, plain, ((e4 = op(e1, e0)) | (~ spl10_16 | ~ spl10_176)), inference(forward_demodulation, [], [f1467, f710])).
fof(f710, plain, ((e0 = op(e4, e1)) | ~ spl10_16), inference(avatar_component_clause, [], [f708])).
fof(f1467, plain, ((e4 = op(e1, op(e4, e1))) | ~ spl10_176), inference(avatar_component_clause, [], [f1465])).
fof(f15141, plain, (spl10_110 | ~ spl10_24 | ~ spl10_182), inference(avatar_split_clause, [], [f15140, f1497, f741, f1102])).
fof(f15140, plain, ((e4 = op(e0, e3)) | (~ spl10_24 | ~ spl10_182)), inference(forward_demodulation, [], [f1499, f743])).
fof(f743, plain, ((e3 = op(e4, e0)) | ~ spl10_24), inference(avatar_component_clause, [], [f741])).
fof(f1499, plain, ((e4 = op(e0, op(e4, e0))) | ~ spl10_182), inference(avatar_component_clause, [], [f1497])).
fof(f15124, plain, (~ spl10_234 | ~ spl10_241), inference(avatar_contradiction_clause, [], [f15123])).
fof(f15123, plain, ($false | (~ spl10_234 | ~ spl10_241)), inference(subsumption_resolution, [], [f15122, f461])).
fof(f461, plain, ~ (e2 = e3), inference(cnf_transformation, [], [f5])).
fof(f15122, plain, ((e2 = e3) | (~ spl10_234 | ~ spl10_241)), inference(forward_demodulation, [], [f1820, f1768])).
fof(f1768, plain, ((e2 = op(op(e0, e0), op(e0, e0))) | ~ spl10_234), inference(avatar_component_clause, [], [f1767])).
fof(f1820, plain, ((e3 = op(op(e0, e0), op(e0, e0))) | ~ spl10_241), inference(avatar_component_clause, [], [f1819])).
fof(f15102, plain, (~ spl10_99 | ~ spl10_81 | spl10_139), inference(avatar_split_clause, [], [f14943, f1278, f981, f1056])).
fof(f1278, plain, (spl10_139 <=> (e3 = op(e1, op(e1, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl10_139])])).
fof(f14943, plain, (~ (e3 = op(e1, e0)) | (~ spl10_81 | spl10_139)), inference(backward_demodulation, [], [f1280, f983])).
fof(f1280, plain, (~ (e3 = op(e1, op(e1, e3))) | spl10_139), inference(avatar_component_clause, [], [f1278])).
fof(f15092, plain, (~ spl10_50 | ~ spl10_26 | spl10_163), inference(avatar_split_clause, [], [f15051, f1391, f750, f850])).
fof(f15051, plain, (~ (e4 = op(e3, e0)) | (~ spl10_26 | spl10_163)), inference(backward_demodulation, [], [f1393, f752])).
fof(f752, plain, ((e0 = op(e3, e4)) | ~ spl10_26), inference(avatar_component_clause, [], [f750])).
fof(f15084, plain, (~ spl10_4 | ~ spl10_24), inference(avatar_split_clause, [], [f15083, f741, f657])).
fof(f657, plain, (spl10_4 <=> (e3 = op(e4, e4))), introduced(avatar_definition, [new_symbols(naming, [spl10_4])])).
fof(f15083, plain, (~ (e3 = op(e4, e4)) | ~ spl10_24), inference(forward_demodulation, [], [f447, f743])).
fof(f447, plain, ~ (op(e4, e0) = op(e4, e4)), inference(cnf_transformation, [], [f4])).
fof(f15079, plain, (spl10_3 | ~ spl10_15 | ~ spl10_173), inference(avatar_split_clause, [], [f15073, f1449, f703, f653])).
fof(f1449, plain, (spl10_173 <=> (e2 = op(e4, op(e4, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl10_173])])).
fof(f15073, plain, ((e2 = op(e4, e4)) | (~ spl10_15 | ~ spl10_173)), inference(backward_demodulation, [], [f1450, f705])).
fof(f1450, plain, ((e2 = op(e4, op(e4, e2))) | ~ spl10_173), inference(avatar_component_clause, [], [f1449])).
fof(f15069, plain, (~ spl10_6 | ~ spl10_16), inference(avatar_split_clause, [], [f15062, f708, f666])).
fof(f15062, plain, (~ (e0 = op(e4, e3)) | ~ spl10_16), inference(backward_demodulation, [], [f449, f710])).
fof(f15059, plain, (~ spl10_19 | ~ spl10_24), inference(avatar_split_clause, [], [f15054, f741, f720])).
fof(f15054, plain, (~ (e3 = op(e4, e1)) | ~ spl10_24), inference(backward_demodulation, [], [f444, f743])).
fof(f444, plain, ~ (op(e4, e0) = op(e4, e1)), inference(cnf_transformation, [], [f4])).
fof(f15052, plain, (spl10_24 | ~ spl10_26 | ~ spl10_160), inference(avatar_split_clause, [], [f15050, f1377, f750, f741])).
fof(f1377, plain, (spl10_160 <=> (e3 = op(e4, op(e3, e4)))), introduced(avatar_definition, [new_symbols(naming, [spl10_160])])).
fof(f15050, plain, ((e3 = op(e4, e0)) | (~ spl10_26 | ~ spl10_160)), inference(backward_demodulation, [], [f1379, f752])).
fof(f1379, plain, ((e3 = op(e4, op(e3, e4))) | ~ spl10_160), inference(avatar_component_clause, [], [f1377])).
fof(f15030, plain, (~ spl10_33 | ~ spl10_39 | spl10_172), inference(avatar_split_clause, [], [f15029, f1443, f804, f779])).
fof(f1443, plain, (spl10_172 <=> (e2 = op(e3, op(e3, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl10_172])])).
fof(f15029, plain, (~ (e2 = op(e3, e3)) | (~ spl10_39 | spl10_172)), inference(backward_demodulation, [], [f1445, f806])).
fof(f1445, plain, (~ (e2 = op(e3, op(e3, e2))) | spl10_172), inference(avatar_component_clause, [], [f1443])).
fof(f15014, plain, (~ spl10_41 | ~ spl10_47 | spl10_184), inference(avatar_split_clause, [], [f15012, f1507, f838, f813])).
fof(f1507, plain, (spl10_184 <=> (e0 = op(e3, op(e3, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl10_184])])).
fof(f15012, plain, (~ (e0 = op(e3, e1)) | (~ spl10_47 | spl10_184)), inference(backward_demodulation, [], [f1509, f840])).
fof(f1509, plain, (~ (e0 = op(e3, op(e3, e0))) | spl10_184), inference(avatar_component_clause, [], [f1507])).
fof(f15013, plain, (~ spl10_27 | ~ spl10_47), inference(avatar_split_clause, [], [f15007, f838, f754])).
fof(f754, plain, (spl10_27 <=> (e1 = op(e3, e4))), introduced(avatar_definition, [new_symbols(naming, [spl10_27])])).
fof(f15007, plain, (~ (e1 = op(e3, e4)) | ~ spl10_47), inference(backward_demodulation, [], [f437, f840])).
fof(f437, plain, ~ (op(e3, e0) = op(e3, e4)), inference(cnf_transformation, [], [f4])).
fof(f14953, plain, (~ spl10_54 | ~ spl10_79), inference(avatar_split_clause, [], [f14948, f972, f867])).
fof(f867, plain, (spl10_54 <=> (e3 = op(e2, e4))), introduced(avatar_definition, [new_symbols(naming, [spl10_54])])).
fof(f14948, plain, (~ (e3 = op(e2, e4)) | ~ spl10_79), inference(backward_demodulation, [], [f398, f974])).
fof(f974, plain, ((e3 = op(e1, e4)) | ~ spl10_79), inference(avatar_component_clause, [], [f972])).
fof(f398, plain, ~ (op(e1, e4) = op(e2, e4)), inference(cnf_transformation, [], [f4])).
fof(f14945, plain, (~ spl10_76 | ~ spl10_81), inference(avatar_split_clause, [], [f14937, f981, f960])).
fof(f14937, plain, (~ (e0 = op(e1, e4)) | ~ spl10_81), inference(backward_demodulation, [], [f423, f983])).
fof(f423, plain, ~ (op(e1, e3) = op(e1, e4)), inference(cnf_transformation, [], [f4])).
fof(f14944, plain, (~ spl10_6 | ~ spl10_81), inference(avatar_split_clause, [], [f14936, f981, f666])).
fof(f14936, plain, (~ (e0 = op(e4, e3)) | ~ spl10_81), inference(backward_demodulation, [], [f390, f983])).
fof(f390, plain, ~ (op(e1, e3) = op(e4, e3)), inference(cnf_transformation, [], [f4])).
fof(f14932, plain, (~ spl10_12 | ~ spl10_87), inference(avatar_split_clause, [], [f14924, f1006, f691])).
fof(f691, plain, (spl10_12 <=> (e1 = op(e4, e2))), introduced(avatar_definition, [new_symbols(naming, [spl10_12])])).
fof(f14924, plain, (~ (e1 = op(e4, e2)) | ~ spl10_87), inference(backward_demodulation, [], [f380, f1008])).
fof(f380, plain, ~ (op(e1, e2) = op(e4, e2)), inference(cnf_transformation, [], [f4])).
fof(f14921, plain, (~ spl10_26 | ~ spl10_95 | ~ spl10_200 | spl10_231), inference(avatar_split_clause, [], [f14916, f1749, f1583, f1039, f750])).
fof(f14916, plain, (~ (e0 = op(e3, e4)) | (~ spl10_95 | ~ spl10_200 | spl10_231)), inference(backward_demodulation, [], [f14793, f1041])).
fof(f14920, plain, (~ spl10_76 | ~ spl10_95 | spl10_237), inference(avatar_split_clause, [], [f14915, f1787, f1039, f960])).
fof(f14915, plain, (~ (e0 = op(e1, e4)) | (~ spl10_95 | spl10_237)), inference(backward_demodulation, [], [f1789, f1041])).
fof(f14919, plain, (spl10_79 | ~ spl10_95 | ~ spl10_200), inference(avatar_split_clause, [], [f14912, f1583, f1039, f972])).
fof(f14912, plain, ((e3 = op(e1, e4)) | (~ spl10_95 | ~ spl10_200)), inference(backward_demodulation, [], [f1584, f1041])).
fof(f14918, plain, (~ spl10_85 | ~ spl10_95), inference(avatar_split_clause, [], [f14909, f1039, f997])).
fof(f14909, plain, (~ (e4 = op(e1, e3)) | ~ spl10_95), inference(backward_demodulation, [], [f419, f1041])).
fof(f419, plain, ~ (op(e1, e1) = op(e1, e3)), inference(cnf_transformation, [], [f4])).
fof(f14902, plain, (~ spl10_73 | ~ spl10_98), inference(avatar_split_clause, [], [f14891, f1052, f947])).
fof(f14891, plain, (~ (e2 = op(e2, e0)) | ~ spl10_98), inference(backward_demodulation, [], [f358, f1054])).
fof(f14890, plain, (spl10_16 | ~ spl10_102 | ~ spl10_135), inference(avatar_split_clause, [], [f14888, f1259, f1069, f708])).
fof(f14888, plain, ((e0 = op(e4, e1)) | (~ spl10_102 | ~ spl10_135)), inference(backward_demodulation, [], [f1261, f1071])).
fof(f1071, plain, ((e1 = op(e0, e4)) | ~ spl10_102), inference(avatar_component_clause, [], [f1069])).
fof(f14889, plain, (~ spl10_27 | ~ spl10_102), inference(avatar_split_clause, [], [f14884, f1069, f754])).
fof(f14884, plain, (~ (e1 = op(e3, e4)) | ~ spl10_102), inference(backward_demodulation, [], [f396, f1071])).
fof(f14856, plain, (spl10_81 | ~ spl10_119 | ~ spl10_152), inference(avatar_split_clause, [], [f14852, f1339, f1140, f981])).
fof(f14852, plain, ((e0 = op(e1, e3)) | (~ spl10_119 | ~ spl10_152)), inference(backward_demodulation, [], [f1341, f1142])).
fof(f14855, plain, (~ spl10_109 | ~ spl10_119), inference(avatar_split_clause, [], [f14849, f1140, f1098])).
fof(f1098, plain, (spl10_109 <=> (e3 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl10_109])])).
fof(f14849, plain, (~ (e3 = op(e0, e3)) | ~ spl10_119), inference(backward_demodulation, [], [f409, f1142])).
fof(f14854, plain, (~ spl10_19 | ~ spl10_119), inference(avatar_split_clause, [], [f14847, f1140, f720])).
fof(f14847, plain, (~ (e3 = op(e4, e1)) | ~ spl10_119), inference(backward_demodulation, [], [f367, f1142])).
fof(f14843, plain, (~ spl10_79 | ~ spl10_125 | spl10_226 | ~ spl10_236), inference(avatar_split_clause, [], [f14839, f1781, f1717, f1165, f972])).
fof(f14839, plain, (~ (e3 = op(e1, e4)) | (~ spl10_125 | spl10_226 | ~ spl10_236)), inference(backward_demodulation, [], [f14786, f1167])).
fof(f1167, plain, ((op(e0, e0) = e4) | ~ spl10_125), inference(avatar_component_clause, [], [f1165])).
fof(f14786, plain, (~ (e3 = op(e1, op(e0, e0))) | (spl10_226 | ~ spl10_236)), inference(backward_demodulation, [], [f1719, f1782])).
fof(f14842, plain, (spl10_102 | ~ spl10_125 | ~ spl10_236), inference(avatar_split_clause, [], [f14838, f1781, f1165, f1069])).
fof(f14838, plain, ((e1 = op(e0, e4)) | (~ spl10_125 | ~ spl10_236)), inference(backward_demodulation, [], [f1782, f1167])).
fof(f14840, plain, (~ spl10_25 | ~ spl10_125), inference(avatar_split_clause, [], [f14831, f1165, f745])).
fof(f14831, plain, (~ (e4 = op(e4, e0)) | ~ spl10_125), inference(backward_demodulation, [], [f357, f1167])).
fof(f14802, plain, (~ spl10_12 | ~ spl10_3 | spl10_205), inference(avatar_split_clause, [], [f14801, f1607, f653, f691])).
fof(f1607, plain, (spl10_205 <=> (e1 = op(e4, op(e4, e4)))), introduced(avatar_definition, [new_symbols(naming, [spl10_205])])).
fof(f14801, plain, (~ (e1 = op(e4, e2)) | (~ spl10_3 | spl10_205)), inference(forward_demodulation, [], [f1609, f655])).
fof(f655, plain, ((e2 = op(e4, e4)) | ~ spl10_3), inference(avatar_component_clause, [], [f653])).
fof(f1609, plain, (~ (e1 = op(e4, op(e4, e4))) | spl10_205), inference(avatar_component_clause, [], [f1607])).
fof(f14791, plain, (spl10_63 | ~ spl10_3 | ~ spl10_233), inference(avatar_split_clause, [], [f14790, f1761, f653, f905])).
fof(f14790, plain, ((e2 = op(e2, e2)) | (~ spl10_3 | ~ spl10_233)), inference(forward_demodulation, [], [f1762, f655])).
fof(f1762, plain, ((e2 = op(op(e4, e4), op(e4, e4))) | ~ spl10_233), inference(avatar_component_clause, [], [f1761])).
fof(f14765, plain, (~ spl10_3 | ~ spl10_89 | ~ spl10_205 | spl10_210), inference(avatar_contradiction_clause, [], [f14764])).
fof(f14764, plain, ($false | (~ spl10_3 | ~ spl10_89 | ~ spl10_205 | spl10_210)), inference(subsumption_resolution, [], [f14763, f1016])).
fof(f14763, plain, (~ (e3 = op(e1, e2)) | (~ spl10_3 | ~ spl10_205 | spl10_210)), inference(backward_demodulation, [], [f14646, f655])).
fof(f14646, plain, (~ (e3 = op(e1, op(e4, e4))) | (~ spl10_205 | spl10_210)), inference(backward_demodulation, [], [f1635, f1608])).
fof(f1608, plain, ((e1 = op(e4, op(e4, e4))) | ~ spl10_205), inference(avatar_component_clause, [], [f1607])).
fof(f1635, plain, (~ (e3 = op(op(e4, op(e4, e4)), op(e4, e4))) | spl10_210), inference(avatar_component_clause, [], [f1633])).
fof(f1633, plain, (spl10_210 <=> (e3 = op(op(e4, op(e4, e4)), op(e4, e4)))), introduced(avatar_definition, [new_symbols(naming, [spl10_210])])).
fof(f14756, plain, (~ spl10_1 | ~ spl10_6), inference(avatar_split_clause, [], [f14750, f666, f645])).
fof(f14750, plain, (~ (e0 = op(e4, e4)) | ~ spl10_6), inference(backward_demodulation, [], [f453, f668])).
fof(f668, plain, ((e0 = op(e4, e3)) | ~ spl10_6), inference(avatar_component_clause, [], [f666])).
fof(f14748, plain, (~ spl10_33 | ~ spl10_40 | spl10_193), inference(avatar_contradiction_clause, [], [f14747])).
fof(f14747, plain, ($false | (~ spl10_33 | ~ spl10_40 | spl10_193)), inference(subsumption_resolution, [], [f14741, f810])).
fof(f14741, plain, (~ (e4 = op(e3, e2)) | (~ spl10_33 | spl10_193)), inference(backward_demodulation, [], [f1554, f781])).
fof(f1554, plain, (~ (e4 = op(e3, op(e3, e3))) | spl10_193), inference(avatar_component_clause, [], [f1552])).
fof(f1552, plain, (spl10_193 <=> (e4 = op(e3, op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl10_193])])).
fof(f14719, plain, (~ spl10_36 | ~ spl10_61), inference(avatar_split_clause, [], [f14707, f897, f792])).
fof(f897, plain, (spl10_61 <=> (e0 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl10_61])])).
fof(f14707, plain, (~ (e0 = op(e3, e2)) | ~ spl10_61), inference(backward_demodulation, [], [f381, f899])).
fof(f899, plain, ((e0 = op(e2, e2)) | ~ spl10_61), inference(avatar_component_clause, [], [f897])).
fof(f14699, plain, (~ spl10_1 | ~ spl10_76), inference(avatar_split_clause, [], [f14692, f960, f645])).
fof(f14692, plain, (~ (e0 = op(e4, e4)) | ~ spl10_76), inference(backward_demodulation, [], [f400, f962])).
fof(f962, plain, ((e0 = op(e1, e4)) | ~ spl10_76), inference(avatar_component_clause, [], [f960])).
fof(f14685, plain, (~ spl10_40 | ~ spl10_93 | ~ spl10_200 | spl10_211), inference(avatar_split_clause, [], [f14681, f1639, f1583, f1031, f808])).
fof(f1639, plain, (spl10_211 <=> (e4 = op(op(e1, op(e1, e1)), op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl10_211])])).
fof(f14681, plain, (~ (e4 = op(e3, e2)) | (~ spl10_93 | ~ spl10_200 | spl10_211)), inference(backward_demodulation, [], [f14648, f1033])).
fof(f14648, plain, (~ (e4 = op(e3, op(e1, e1))) | (~ spl10_200 | spl10_211)), inference(backward_demodulation, [], [f1641, f1584])).
fof(f1641, plain, (~ (e4 = op(op(e1, op(e1, e1)), op(e1, e1))) | spl10_211), inference(avatar_component_clause, [], [f1639])).
fof(f14684, plain, (spl10_61 | ~ spl10_93 | ~ spl10_198), inference(avatar_split_clause, [], [f14677, f1575, f1031, f897])).
fof(f1575, plain, (spl10_198 <=> (e0 = op(op(e1, e1), op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl10_198])])).
fof(f14677, plain, ((e0 = op(e2, e2)) | (~ spl10_93 | ~ spl10_198)), inference(backward_demodulation, [], [f1576, f1033])).
fof(f1576, plain, ((e0 = op(op(e1, e1), op(e1, e1))) | ~ spl10_198), inference(avatar_component_clause, [], [f1575])).
fof(f14664, plain, (~ spl10_169 | ~ spl10_188), inference(avatar_contradiction_clause, [], [f14663])).
fof(f14663, plain, ($false | (~ spl10_169 | ~ spl10_188)), inference(subsumption_resolution, [], [f14662, f461])).
fof(f14662, plain, ((e2 = e3) | (~ spl10_169 | ~ spl10_188)), inference(backward_demodulation, [], [f1530, f1431])).
fof(f1431, plain, ((e2 = op(e2, op(e2, e2))) | ~ spl10_169), inference(avatar_component_clause, [], [f1429])).
fof(f1530, plain, ((e3 = op(e2, op(e2, e2))) | ~ spl10_188), inference(avatar_component_clause, [], [f1529])).
fof(f1529, plain, (spl10_188 <=> (e3 = op(e2, op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl10_188])])).
fof(f14661, plain, (spl10_70 | ~ spl10_12 | ~ spl10_171), inference(avatar_split_clause, [], [f14660, f1437, f691, f934])).
fof(f1437, plain, (spl10_171 <=> (e4 = op(e2, op(e4, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl10_171])])).
fof(f14660, plain, ((e4 = op(e2, e1)) | (~ spl10_12 | ~ spl10_171)), inference(forward_demodulation, [], [f1439, f693])).
fof(f693, plain, ((e1 = op(e4, e2)) | ~ spl10_12), inference(avatar_component_clause, [], [f691])).
fof(f1439, plain, ((e4 = op(e2, op(e4, e2))) | ~ spl10_171), inference(avatar_component_clause, [], [f1437])).
fof(f14659, plain, (spl10_85 | ~ spl10_19 | ~ spl10_176), inference(avatar_split_clause, [], [f14658, f1465, f720, f997])).
fof(f14658, plain, ((e4 = op(e1, e3)) | (~ spl10_19 | ~ spl10_176)), inference(forward_demodulation, [], [f1467, f722])).
fof(f14655, plain, (~ spl10_1 | ~ spl10_25 | spl10_185), inference(avatar_split_clause, [], [f14654, f1513, f745, f645])).
fof(f1513, plain, (spl10_185 <=> (e0 = op(e4, op(e4, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl10_185])])).
fof(f14654, plain, (~ (e0 = op(e4, e4)) | (~ spl10_25 | spl10_185)), inference(forward_demodulation, [], [f1515, f747])).
fof(f747, plain, ((e4 = op(e4, e0)) | ~ spl10_25), inference(avatar_component_clause, [], [f745])).
fof(f1515, plain, (~ (e0 = op(e4, op(e4, e0))) | spl10_185), inference(avatar_component_clause, [], [f1513])).
fof(f14633, plain, (~ spl10_69 | ~ spl10_57 | spl10_167), inference(avatar_split_clause, [], [f14536, f1414, f880, f930])).
fof(f1414, plain, (spl10_167 <=> (e3 = op(e2, op(e2, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl10_167])])).
fof(f14536, plain, (~ (e3 = op(e2, e1)) | (~ spl10_57 | spl10_167)), inference(backward_demodulation, [], [f1416, f882])).
fof(f1416, plain, (~ (e3 = op(e2, op(e2, e3))) | spl10_167), inference(avatar_component_clause, [], [f1414])).
fof(f14627, plain, (~ spl10_45 | ~ spl10_27 | spl10_163), inference(avatar_split_clause, [], [f14575, f1391, f754, f829])).
fof(f14575, plain, (~ (e4 = op(e3, e1)) | (~ spl10_27 | spl10_163)), inference(backward_demodulation, [], [f1393, f756])).
fof(f756, plain, ((e1 = op(e3, e4)) | ~ spl10_27), inference(avatar_component_clause, [], [f754])).
fof(f14617, plain, (~ spl10_1 | ~ spl10_25 | spl10_161), inference(avatar_contradiction_clause, [], [f14616])).
fof(f14616, plain, ($false | (~ spl10_1 | ~ spl10_25 | spl10_161)), inference(subsumption_resolution, [], [f14607, f747])).
fof(f14607, plain, (~ (e4 = op(e4, e0)) | (~ spl10_1 | spl10_161)), inference(backward_demodulation, [], [f1382, f647])).
fof(f1382, plain, (~ (e4 = op(e4, op(e4, e4))) | spl10_161), inference(avatar_component_clause, [], [f1381])).
fof(f1381, plain, (spl10_161 <=> (e4 = op(e4, op(e4, e4)))), introduced(avatar_definition, [new_symbols(naming, [spl10_161])])).
fof(f14594, plain, (~ spl10_2 | ~ spl10_12), inference(avatar_split_clause, [], [f14588, f691, f649])).
fof(f14588, plain, (~ (e1 = op(e4, e4)) | ~ spl10_12), inference(backward_demodulation, [], [f452, f693])).
fof(f14585, plain, (~ spl10_14 | ~ spl10_19), inference(avatar_split_clause, [], [f14579, f720, f699])).
fof(f699, plain, (spl10_14 <=> (e3 = op(e4, e2))), introduced(avatar_definition, [new_symbols(naming, [spl10_14])])).
fof(f14579, plain, (~ (e3 = op(e4, e2)) | ~ spl10_19), inference(backward_demodulation, [], [f448, f722])).
fof(f448, plain, ~ (op(e4, e1) = op(e4, e2)), inference(cnf_transformation, [], [f4])).
fof(f14577, plain, (spl10_19 | ~ spl10_27 | ~ spl10_160), inference(avatar_split_clause, [], [f14574, f1377, f754, f720])).
fof(f14574, plain, ((e3 = op(e4, e1)) | (~ spl10_27 | ~ spl10_160)), inference(backward_demodulation, [], [f1379, f756])).
fof(f14576, plain, (~ spl10_2 | ~ spl10_27), inference(avatar_split_clause, [], [f14570, f754, f649])).
fof(f14570, plain, (~ (e1 = op(e4, e4)) | ~ spl10_27), inference(backward_demodulation, [], [f403, f756])).
fof(f403, plain, ~ (op(e3, e4) = op(e4, e4)), inference(cnf_transformation, [], [f4])).
fof(f14566, plain, (spl10_1 | ~ spl10_35 | ~ spl10_190), inference(avatar_split_clause, [], [f14559, f1539, f787, f645])).
fof(f14559, plain, ((e0 = op(e4, e4)) | (~ spl10_35 | ~ spl10_190)), inference(backward_demodulation, [], [f1540, f789])).
fof(f1540, plain, ((e0 = op(op(e3, e3), op(e3, e3))) | ~ spl10_190), inference(avatar_component_clause, [], [f1539])).
fof(f14555, plain, (~ spl10_28 | ~ spl10_43), inference(avatar_split_clause, [], [f14549, f821, f758])).
fof(f14549, plain, (~ (e2 = op(e3, e4)) | ~ spl10_43), inference(backward_demodulation, [], [f440, f823])).
fof(f14554, plain, (~ spl10_18 | ~ spl10_43), inference(avatar_split_clause, [], [f14547, f821, f716])).
fof(f716, plain, (spl10_18 <=> (e2 = op(e4, e1))), introduced(avatar_definition, [new_symbols(naming, [spl10_18])])).
fof(f14547, plain, (~ (e2 = op(e4, e1)) | ~ spl10_43), inference(backward_demodulation, [], [f373, f823])).
fof(f14545, plain, (spl10_8 | ~ spl10_54 | ~ spl10_159), inference(avatar_split_clause, [], [f14543, f1373, f867, f674])).
fof(f1373, plain, (spl10_159 <=> (e2 = op(e4, op(e2, e4)))), introduced(avatar_definition, [new_symbols(naming, [spl10_159])])).
fof(f14543, plain, ((e2 = op(e4, e3)) | (~ spl10_54 | ~ spl10_159)), inference(backward_demodulation, [], [f1375, f869])).
fof(f869, plain, ((e3 = op(e2, e4)) | ~ spl10_54), inference(avatar_component_clause, [], [f867])).
fof(f1375, plain, ((e2 = op(e4, op(e2, e4))) | ~ spl10_159), inference(avatar_component_clause, [], [f1373])).
fof(f14528, plain, (~ spl10_27 | ~ spl10_65 | spl10_187 | ~ spl10_188), inference(avatar_split_clause, [], [f14522, f1529, f1525, f913, f754])).
fof(f14522, plain, (~ (e1 = op(e3, e4)) | (~ spl10_65 | spl10_187 | ~ spl10_188)), inference(backward_demodulation, [], [f14174, f915])).
fof(f14174, plain, (~ (e1 = op(e3, op(e2, e2))) | (spl10_187 | ~ spl10_188)), inference(backward_demodulation, [], [f1527, f1530])).
fof(f14527, plain, (~ spl10_51 | ~ spl10_65 | spl10_228), inference(avatar_split_clause, [], [f14520, f1727, f913, f855])).
fof(f855, plain, (spl10_51 <=> (e0 = op(e2, e4))), introduced(avatar_definition, [new_symbols(naming, [spl10_51])])).
fof(f14520, plain, (~ (e0 = op(e2, e4)) | (~ spl10_65 | spl10_228)), inference(backward_demodulation, [], [f1729, f915])).
fof(f1729, plain, (~ (e0 = op(e2, op(e2, e2))) | spl10_228), inference(avatar_component_clause, [], [f1727])).
fof(f14524, plain, (~ spl10_60 | ~ spl10_65), inference(avatar_split_clause, [], [f14515, f913, f892])).
fof(f14515, plain, (~ (e4 = op(e2, e3)) | ~ spl10_65), inference(backward_demodulation, [], [f431, f915])).
fof(f431, plain, ~ (op(e2, e2) = op(e2, e3)), inference(cnf_transformation, [], [f4])).
fof(f14512, plain, (~ spl10_51 | ~ spl10_66), inference(avatar_split_clause, [], [f14505, f918, f855])).
fof(f14505, plain, (~ (e0 = op(e2, e4)) | ~ spl10_66), inference(backward_demodulation, [], [f430, f920])).
fof(f430, plain, ~ (op(e2, e1) = op(e2, e4)), inference(cnf_transformation, [], [f4])).
fof(f14499, plain, (spl10_12 | ~ spl10_78 | ~ spl10_132), inference(avatar_split_clause, [], [f14495, f1245, f968, f691])).
fof(f1245, plain, (spl10_132 <=> (e1 = op(e4, op(e1, e4)))), introduced(avatar_definition, [new_symbols(naming, [spl10_132])])).
fof(f14495, plain, ((e1 = op(e4, e2)) | (~ spl10_78 | ~ spl10_132)), inference(backward_demodulation, [], [f1247, f970])).
fof(f1247, plain, ((e1 = op(e4, op(e1, e4))) | ~ spl10_132), inference(avatar_component_clause, [], [f1245])).
fof(f14498, plain, (~ spl10_28 | ~ spl10_78), inference(avatar_split_clause, [], [f14491, f968, f758])).
fof(f14491, plain, (~ (e2 = op(e3, e4)) | ~ spl10_78), inference(backward_demodulation, [], [f399, f970])).
fof(f399, plain, ~ (op(e1, e4) = op(e3, e4)), inference(cnf_transformation, [], [f4])).
fof(f14477, plain, (~ spl10_83 | ~ spl10_89 | spl10_145), inference(avatar_split_clause, [], [f14473, f1306, f1014, f989])).
fof(f1306, plain, (spl10_145 <=> (e2 = op(e1, op(e1, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl10_145])])).
fof(f14473, plain, (~ (e2 = op(e1, e3)) | (~ spl10_89 | spl10_145)), inference(backward_demodulation, [], [f1308, f1016])).
fof(f1308, plain, (~ (e2 = op(e1, op(e1, e2))) | spl10_145), inference(avatar_component_clause, [], [f1306])).
fof(f14475, plain, (~ spl10_79 | ~ spl10_89), inference(avatar_split_clause, [], [f14470, f1014, f972])).
fof(f14470, plain, (~ (e3 = op(e1, e4)) | ~ spl10_89), inference(backward_demodulation, [], [f422, f1016])).
fof(f14474, plain, (~ spl10_14 | ~ spl10_89), inference(avatar_split_clause, [], [f14468, f1014, f699])).
fof(f14468, plain, (~ (e3 = op(e4, e2)) | ~ spl10_89), inference(backward_demodulation, [], [f380, f1016])).
fof(f14466, plain, (~ spl10_78 | ~ spl10_95 | spl10_207), inference(avatar_split_clause, [], [f14459, f1617, f1039, f968])).
fof(f14459, plain, (~ (e2 = op(e1, e4)) | (~ spl10_95 | spl10_207)), inference(backward_demodulation, [], [f1619, f1041])).
fof(f1619, plain, (~ (e2 = op(e1, op(e1, e1))) | spl10_207), inference(avatar_component_clause, [], [f1617])).
fof(f14465, plain, (~ spl10_79 | ~ spl10_95 | spl10_200), inference(avatar_split_clause, [], [f14457, f1583, f1039, f972])).
fof(f14457, plain, (~ (e3 = op(e1, e4)) | (~ spl10_95 | spl10_200)), inference(backward_demodulation, [], [f1585, f1041])).
fof(f1585, plain, (~ (e3 = op(e1, op(e1, e1))) | spl10_200), inference(avatar_component_clause, [], [f1583])).
fof(f14464, plain, (spl10_1 | ~ spl10_95 | ~ spl10_198), inference(avatar_split_clause, [], [f14456, f1575, f1039, f645])).
fof(f14456, plain, ((e0 = op(e4, e4)) | (~ spl10_95 | ~ spl10_198)), inference(backward_demodulation, [], [f1576, f1041])).
fof(f14463, plain, (~ spl10_90 | ~ spl10_95), inference(avatar_split_clause, [], [f14453, f1039, f1018])).
fof(f14453, plain, (~ (e4 = op(e1, e2)) | ~ spl10_95), inference(backward_demodulation, [], [f418, f1041])).
fof(f14462, plain, (~ spl10_45 | ~ spl10_95), inference(avatar_split_clause, [], [f14451, f1039, f829])).
fof(f14451, plain, (~ (e4 = op(e3, e1)) | ~ spl10_95), inference(backward_demodulation, [], [f369, f1041])).
fof(f14449, plain, (~ spl10_91 | ~ spl10_117 | spl10_152), inference(avatar_split_clause, [], [f14448, f1339, f1132, f1023])).
fof(f1023, plain, (spl10_91 <=> (e0 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl10_91])])).
fof(f1132, plain, (spl10_117 <=> (e1 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl10_117])])).
fof(f14448, plain, (~ (e0 = op(e1, e1)) | (~ spl10_117 | spl10_152)), inference(forward_demodulation, [], [f1340, f1134])).
fof(f1134, plain, ((e1 = op(e0, e1)) | ~ spl10_117), inference(avatar_component_clause, [], [f1132])).
fof(f14447, plain, (~ spl10_91 | ~ spl10_97 | spl10_156), inference(avatar_split_clause, [], [f14446, f1358, f1048, f1023])).
fof(f1048, plain, (spl10_97 <=> (e1 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl10_97])])).
fof(f14446, plain, (~ (e0 = op(e1, e1)) | (~ spl10_97 | spl10_156)), inference(forward_demodulation, [], [f1360, f1050])).
fof(f1050, plain, ((e1 = op(e1, e0)) | ~ spl10_97), inference(avatar_component_clause, [], [f1048])).
fof(f14431, plain, (~ spl10_92 | ~ spl10_97), inference(avatar_split_clause, [], [f14430, f1048, f1027])).
fof(f1027, plain, (spl10_92 <=> (e1 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl10_92])])).
fof(f14430, plain, (~ (e1 = op(e1, e1)) | ~ spl10_97), inference(forward_demodulation, [], [f414, f1050])).
fof(f14429, plain, (~ spl10_92 | ~ spl10_117), inference(avatar_split_clause, [], [f14428, f1132, f1027])).
fof(f14428, plain, (~ (e1 = op(e1, e1)) | ~ spl10_117), inference(forward_demodulation, [], [f364, f1134])).
fof(f14423, plain, (~ spl10_84 | ~ spl10_109), inference(avatar_split_clause, [], [f14422, f1098, f993])).
fof(f993, plain, (spl10_84 <=> (e3 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl10_84])])).
fof(f14422, plain, (~ (e3 = op(e1, e3)) | ~ spl10_109), inference(forward_demodulation, [], [f384, f1100])).
fof(f1100, plain, ((e3 = op(e0, e3)) | ~ spl10_109), inference(avatar_component_clause, [], [f1098])).
fof(f14413, plain, (~ spl10_1 | ~ spl10_105 | spl10_135), inference(avatar_split_clause, [], [f14268, f1259, f1081, f645])).
fof(f14268, plain, (~ (e0 = op(e4, e4)) | (~ spl10_105 | spl10_135)), inference(backward_demodulation, [], [f1260, f1083])).
fof(f1083, plain, ((e4 = op(e0, e4)) | ~ spl10_105), inference(avatar_component_clause, [], [f1081])).
fof(f1260, plain, (~ (e0 = op(e4, op(e0, e4))) | spl10_135), inference(avatar_component_clause, [], [f1259])).
fof(f14392, plain, (~ spl10_9 | ~ spl10_14), inference(avatar_split_clause, [], [f14387, f699, f678])).
fof(f14387, plain, (~ (e3 = op(e4, e3)) | ~ spl10_14), inference(backward_demodulation, [], [f451, f701])).
fof(f701, plain, ((e3 = op(e4, e2)) | ~ spl10_14), inference(avatar_component_clause, [], [f699])).
fof(f14386, plain, (~ spl10_13 | ~ spl10_18), inference(avatar_split_clause, [], [f14379, f716, f695])).
fof(f14379, plain, (~ (e2 = op(e4, e2)) | ~ spl10_18), inference(backward_demodulation, [], [f448, f718])).
fof(f718, plain, ((e2 = op(e4, e1)) | ~ spl10_18), inference(avatar_component_clause, [], [f716])).
fof(f14378, plain, (~ spl10_25 | ~ spl10_105 | spl10_182), inference(avatar_contradiction_clause, [], [f14377])).
fof(f14377, plain, ($false | (~ spl10_25 | ~ spl10_105 | spl10_182)), inference(subsumption_resolution, [], [f14375, f1083])).
fof(f14375, plain, (~ (e4 = op(e0, e4)) | (~ spl10_25 | spl10_182)), inference(backward_demodulation, [], [f1498, f747])).
fof(f14376, plain, (~ spl10_5 | ~ spl10_25), inference(avatar_split_clause, [], [f14374, f745, f661])).
fof(f14374, plain, (~ (e4 = op(e4, e4)) | ~ spl10_25), inference(backward_demodulation, [], [f447, f747])).
fof(f14369, plain, (spl10_14 | ~ spl10_28 | ~ spl10_160), inference(avatar_split_clause, [], [f14367, f1377, f758, f699])).
fof(f14367, plain, ((e3 = op(e4, e2)) | (~ spl10_28 | ~ spl10_160)), inference(backward_demodulation, [], [f1379, f760])).
fof(f760, plain, ((e2 = op(e3, e4)) | ~ spl10_28), inference(avatar_component_clause, [], [f758])).
fof(f14362, plain, (~ spl10_18 | ~ spl10_32 | ~ spl10_193 | spl10_202), inference(avatar_split_clause, [], [f14359, f1593, f1552, f775, f716])).
fof(f14359, plain, (~ (e2 = op(e4, e1)) | (~ spl10_32 | ~ spl10_193 | spl10_202)), inference(backward_demodulation, [], [f14167, f777])).
fof(f14167, plain, (~ (e2 = op(e4, op(e3, e3))) | (~ spl10_193 | spl10_202)), inference(backward_demodulation, [], [f1595, f1553])).
fof(f1553, plain, ((e4 = op(e3, op(e3, e3))) | ~ spl10_193), inference(avatar_component_clause, [], [f1552])).
fof(f14346, plain, (~ spl10_31 | ~ spl10_36), inference(avatar_split_clause, [], [f14339, f792, f771])).
fof(f771, plain, (spl10_31 <=> (e0 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl10_31])])).
fof(f14339, plain, (~ (e0 = op(e3, e3)) | ~ spl10_36), inference(backward_demodulation, [], [f441, f794])).
fof(f14337, plain, (~ spl10_29 | ~ spl10_49), inference(avatar_split_clause, [], [f14334, f846, f762])).
fof(f846, plain, (spl10_49 <=> (e3 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl10_49])])).
fof(f14334, plain, (~ (e3 = op(e3, e4)) | ~ spl10_49), inference(backward_demodulation, [], [f437, f848])).
fof(f848, plain, ((e3 = op(e3, e0)) | ~ spl10_49), inference(avatar_component_clause, [], [f846])).
fof(f14301, plain, (~ spl10_61 | ~ spl10_73 | spl10_183), inference(avatar_split_clause, [], [f14298, f1502, f947, f897])).
fof(f1502, plain, (spl10_183 <=> (e0 = op(e2, op(e2, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl10_183])])).
fof(f14298, plain, (~ (e0 = op(e2, e2)) | (~ spl10_73 | spl10_183)), inference(backward_demodulation, [], [f1504, f949])).
fof(f949, plain, ((e2 = op(e2, e0)) | ~ spl10_73), inference(avatar_component_clause, [], [f947])).
fof(f1504, plain, (~ (e0 = op(e2, op(e2, e0))) | spl10_183), inference(avatar_component_clause, [], [f1502])).
fof(f14300, plain, (~ spl10_53 | ~ spl10_73), inference(avatar_split_clause, [], [f14294, f947, f863])).
fof(f14294, plain, (~ (e2 = op(e2, e4)) | ~ spl10_73), inference(backward_demodulation, [], [f427, f949])).
fof(f427, plain, ~ (op(e2, e0) = op(e2, e4)), inference(cnf_transformation, [], [f4])).
fof(f14289, plain, (~ spl10_29 | ~ spl10_79), inference(avatar_split_clause, [], [f14284, f972, f762])).
fof(f14284, plain, (~ (e3 = op(e3, e4)) | ~ spl10_79), inference(backward_demodulation, [], [f399, f974])).
fof(f14281, plain, (~ spl10_77 | ~ spl10_97), inference(avatar_split_clause, [], [f14274, f1048, f964])).
fof(f14274, plain, (~ (e1 = op(e1, e4)) | ~ spl10_97), inference(backward_demodulation, [], [f417, f1050])).
fof(f417, plain, ~ (op(e1, e0) = op(e1, e4)), inference(cnf_transformation, [], [f4])).
fof(f14270, plain, (~ spl10_5 | ~ spl10_105), inference(avatar_split_clause, [], [f14267, f1081, f661])).
fof(f14267, plain, (~ (e4 = op(e4, e4)) | ~ spl10_105), inference(backward_demodulation, [], [f397, f1083])).
fof(f397, plain, ~ (op(e0, e4) = op(e4, e4)), inference(cnf_transformation, [], [f4])).
fof(f14263, plain, (~ spl10_31 | ~ spl10_109 | spl10_141), inference(avatar_split_clause, [], [f14260, f1287, f1098, f771])).
fof(f14260, plain, (~ (e0 = op(e3, e3)) | (~ spl10_109 | spl10_141)), inference(backward_demodulation, [], [f1288, f1100])).
fof(f14262, plain, (~ spl10_9 | ~ spl10_109), inference(avatar_split_clause, [], [f14257, f1098, f678])).
fof(f14257, plain, (~ (e3 = op(e4, e3)) | ~ spl10_109), inference(backward_demodulation, [], [f387, f1100])).
fof(f387, plain, ~ (op(e0, e3) = op(e4, e3)), inference(cnf_transformation, [], [f4])).
fof(f14254, plain, (~ spl10_13 | ~ spl10_113), inference(avatar_split_clause, [], [f14247, f1115, f695])).
fof(f14247, plain, (~ (e2 = op(e4, e2)) | ~ spl10_113), inference(backward_demodulation, [], [f377, f1117])).
fof(f1117, plain, ((e2 = op(e0, e2)) | ~ spl10_113), inference(avatar_component_clause, [], [f1115])).
fof(f377, plain, ~ (op(e0, e2) = op(e4, e2)), inference(cnf_transformation, [], [f4])).
fof(f14244, plain, (~ spl10_107 | ~ spl10_117), inference(avatar_split_clause, [], [f14236, f1132, f1090])).
fof(f14236, plain, (~ (e1 = op(e0, e3)) | ~ spl10_117), inference(backward_demodulation, [], [f409, f1134])).
fof(f14243, plain, (~ spl10_17 | ~ spl10_117), inference(avatar_split_clause, [], [f14234, f1132, f712])).
fof(f14234, plain, (~ (e1 = op(e4, e1)) | ~ spl10_117), inference(backward_demodulation, [], [f367, f1134])).
fof(f14225, plain, (~ spl10_101 | ~ spl10_121), inference(avatar_split_clause, [], [f14212, f1149, f1065])).
fof(f1149, plain, (spl10_121 <=> (e0 = op(e0, e0))), introduced(avatar_definition, [new_symbols(naming, [spl10_121])])).
fof(f14212, plain, (~ (e0 = op(e0, e4)) | ~ spl10_121), inference(backward_demodulation, [], [f407, f1151])).
fof(f1151, plain, ((e0 = op(e0, e0)) | ~ spl10_121), inference(avatar_component_clause, [], [f1149])).
fof(f14224, plain, (~ spl10_21 | ~ spl10_121), inference(avatar_split_clause, [], [f14208, f1149, f729])).
fof(f14208, plain, (~ (e0 = op(e4, e0)) | ~ spl10_121), inference(backward_demodulation, [], [f357, f1151])).
fof(f14204, plain, (spl10_25 | ~ spl10_126), inference(avatar_split_clause, [], [f14190, f1170, f745])).
fof(f1170, plain, (spl10_126 <=> (e0 = unit)), introduced(avatar_definition, [new_symbols(naming, [spl10_126])])).
fof(f14190, plain, ((e4 = op(e4, e0)) | ~ spl10_126), inference(backward_demodulation, [], [f302, f1172])).
fof(f1172, plain, ((e0 = unit) | ~ spl10_126), inference(avatar_component_clause, [], [f1170])).
fof(f14203, plain, (spl10_105 | ~ spl10_126), inference(avatar_split_clause, [], [f14189, f1170, f1081])).
fof(f14189, plain, ((e4 = op(e0, e4)) | ~ spl10_126), inference(backward_demodulation, [], [f301, f1172])).
fof(f14202, plain, (spl10_49 | ~ spl10_126), inference(avatar_split_clause, [], [f14188, f1170, f846])).
fof(f14188, plain, ((e3 = op(e3, e0)) | ~ spl10_126), inference(backward_demodulation, [], [f300, f1172])).
fof(f14201, plain, (spl10_109 | ~ spl10_126), inference(avatar_split_clause, [], [f14187, f1170, f1098])).
fof(f14187, plain, ((e3 = op(e0, e3)) | ~ spl10_126), inference(backward_demodulation, [], [f299, f1172])).
fof(f14200, plain, (spl10_73 | ~ spl10_126), inference(avatar_split_clause, [], [f14186, f1170, f947])).
fof(f14186, plain, ((e2 = op(e2, e0)) | ~ spl10_126), inference(backward_demodulation, [], [f298, f1172])).
fof(f14199, plain, (spl10_113 | ~ spl10_126), inference(avatar_split_clause, [], [f14185, f1170, f1115])).
fof(f14185, plain, ((e2 = op(e0, e2)) | ~ spl10_126), inference(backward_demodulation, [], [f297, f1172])).
fof(f14198, plain, (spl10_97 | ~ spl10_126), inference(avatar_split_clause, [], [f14184, f1170, f1048])).
fof(f14184, plain, ((e1 = op(e1, e0)) | ~ spl10_126), inference(backward_demodulation, [], [f296, f1172])).
fof(f14197, plain, (spl10_117 | ~ spl10_126), inference(avatar_split_clause, [], [f14183, f1170, f1132])).
fof(f14183, plain, ((e1 = op(e0, e1)) | ~ spl10_126), inference(backward_demodulation, [], [f295, f1172])).
fof(f14196, plain, (spl10_121 | ~ spl10_126), inference(avatar_split_clause, [], [f14182, f1170, f1149])).
fof(f14182, plain, ((e0 = op(e0, e0)) | ~ spl10_126), inference(backward_demodulation, [], [f294, f1172])).
fof(f14177, plain, (spl10_28 | ~ spl10_60 | ~ spl10_164), inference(avatar_split_clause, [], [f14176, f1401, f892, f758])).
fof(f14176, plain, ((e2 = op(e3, e4)) | (~ spl10_60 | ~ spl10_164)), inference(forward_demodulation, [], [f1403, f894])).
fof(f894, plain, ((e4 = op(e2, e3)) | ~ spl10_60), inference(avatar_component_clause, [], [f892])).
fof(f14163, plain, (spl10_121 | ~ spl10_91 | ~ spl10_198), inference(avatar_split_clause, [], [f14162, f1575, f1023, f1149])).
fof(f14162, plain, ((e0 = op(e0, e0)) | (~ spl10_91 | ~ spl10_198)), inference(forward_demodulation, [], [f1576, f1025])).
fof(f1025, plain, ((e0 = op(e1, e1)) | ~ spl10_91), inference(avatar_component_clause, [], [f1023])).
fof(f14161, plain, (~ spl10_99 | ~ spl10_91 | spl10_200), inference(avatar_split_clause, [], [f14160, f1583, f1023, f1056])).
fof(f14160, plain, (~ (e3 = op(e1, e0)) | (~ spl10_91 | spl10_200)), inference(forward_demodulation, [], [f1585, f1025])).
fof(f14152, plain, (~ spl10_78 | ~ spl10_83), inference(avatar_split_clause, [], [f14151, f989, f968])).
fof(f14151, plain, (~ (e2 = op(e1, e4)) | ~ spl10_83), inference(forward_demodulation, [], [f423, f991])).
fof(f14150, plain, (~ spl10_80 | ~ spl10_90), inference(avatar_split_clause, [], [f14149, f1018, f976])).
fof(f976, plain, (spl10_80 <=> (e4 = op(e1, e4))), introduced(avatar_definition, [new_symbols(naming, [spl10_80])])).
fof(f14149, plain, (~ (e4 = op(e1, e4)) | ~ spl10_90), inference(forward_demodulation, [], [f422, f1020])).
fof(f14148, plain, (~ spl10_78 | ~ spl10_90 | spl10_145), inference(avatar_split_clause, [], [f14071, f1306, f1018, f968])).
fof(f14071, plain, (~ (e2 = op(e1, e4)) | (~ spl10_90 | spl10_145)), inference(backward_demodulation, [], [f1308, f1020])).
fof(f14139, plain, (~ spl10_55 | ~ spl10_60), inference(avatar_split_clause, [], [f14138, f892, f871])).
fof(f14138, plain, (~ (e4 = op(e2, e4)) | ~ spl10_60), inference(forward_demodulation, [], [f433, f894])).
fof(f433, plain, ~ (op(e2, e3) = op(e2, e4)), inference(cnf_transformation, [], [f4])).
fof(f14137, plain, (~ spl10_54 | ~ spl10_69), inference(avatar_split_clause, [], [f14136, f930, f867])).
fof(f14136, plain, (~ (e3 = op(e2, e4)) | ~ spl10_69), inference(forward_demodulation, [], [f430, f932])).
fof(f14135, plain, (~ spl10_54 | ~ spl10_60 | spl10_167), inference(avatar_split_clause, [], [f14103, f1414, f892, f867])).
fof(f14103, plain, (~ (e3 = op(e2, e4)) | (~ spl10_60 | spl10_167)), inference(backward_demodulation, [], [f1416, f894])).
fof(f14131, plain, (~ spl10_40 | ~ spl10_90), inference(avatar_split_clause, [], [f14130, f1018, f808])).
fof(f14130, plain, (~ (e4 = op(e3, e2)) | ~ spl10_90), inference(forward_demodulation, [], [f379, f1020])).
fof(f14127, plain, (~ spl10_35 | ~ spl10_60), inference(avatar_split_clause, [], [f14126, f892, f787])).
fof(f14126, plain, (~ (e4 = op(e3, e3)) | ~ spl10_60), inference(forward_demodulation, [], [f391, f894])).
fof(f14117, plain, (~ spl10_15 | ~ spl10_90), inference(avatar_split_clause, [], [f14116, f1018, f703])).
fof(f14116, plain, (~ (e4 = op(e4, e2)) | ~ spl10_90), inference(forward_demodulation, [], [f380, f1020])).
fof(f14115, plain, (~ spl10_10 | ~ spl10_60), inference(avatar_split_clause, [], [f14114, f892, f682])).
fof(f682, plain, (spl10_10 <=> (e4 = op(e4, e3))), introduced(avatar_definition, [new_symbols(naming, [spl10_10])])).
fof(f14114, plain, (~ (e4 = op(e4, e3)) | ~ spl10_60), inference(forward_demodulation, [], [f392, f894])).
fof(f14101, plain, (~ spl10_57 | ~ spl10_69 | spl10_177), inference(avatar_split_clause, [], [f14100, f1470, f930, f880])).
fof(f1470, plain, (spl10_177 <=> (e1 = op(e2, op(e2, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl10_177])])).
fof(f14100, plain, (~ (e1 = op(e2, e3)) | (~ spl10_69 | spl10_177)), inference(backward_demodulation, [], [f1472, f932])).
fof(f1472, plain, (~ (e1 = op(e2, op(e2, e1))) | spl10_177), inference(avatar_component_clause, [], [f1470])).
fof(f14090, plain, (~ spl10_57 | ~ spl10_72), inference(avatar_split_clause, [], [f14081, f943, f880])).
fof(f14081, plain, (~ (e1 = op(e2, e3)) | ~ spl10_72), inference(backward_demodulation, [], [f426, f945])).
fof(f14072, plain, (~ spl10_85 | ~ spl10_90), inference(avatar_split_clause, [], [f14069, f1018, f997])).
fof(f14069, plain, (~ (e4 = op(e1, e3)) | ~ spl10_90), inference(backward_demodulation, [], [f421, f1020])).
fof(f421, plain, ~ (op(e1, e2) = op(e1, e3)), inference(cnf_transformation, [], [f4])).
fof(f14068, plain, (spl10_48 | ~ spl10_91 | ~ spl10_99 | ~ spl10_199), inference(avatar_split_clause, [], [f14065, f1579, f1056, f1023, f842])).
fof(f14065, plain, ((e2 = op(e3, e0)) | (~ spl10_91 | ~ spl10_99 | ~ spl10_199)), inference(backward_demodulation, [], [f13995, f1058])).
fof(f13995, plain, ((e2 = op(op(e1, e0), e0)) | (~ spl10_91 | ~ spl10_199)), inference(forward_demodulation, [], [f1580, f1025])).
fof(f1580, plain, ((e2 = op(op(e1, op(e1, e1)), op(e1, e1))) | ~ spl10_199), inference(avatar_component_clause, [], [f1579])).
fof(f14067, plain, (~ spl10_89 | ~ spl10_99), inference(avatar_split_clause, [], [f14060, f1056, f1014])).
fof(f14060, plain, (~ (e3 = op(e1, e2)) | ~ spl10_99), inference(backward_demodulation, [], [f415, f1058])).
fof(f14066, plain, (~ spl10_74 | ~ spl10_99), inference(avatar_split_clause, [], [f14058, f1056, f951])).
fof(f14058, plain, (~ (e3 = op(e2, e0)) | ~ spl10_99), inference(backward_demodulation, [], [f358, f1058])).
fof(f14056, plain, (~ spl10_57 | ~ spl10_107), inference(avatar_split_clause, [], [f14050, f1090, f880])).
fof(f14050, plain, (~ (e1 = op(e2, e3)) | ~ spl10_107), inference(backward_demodulation, [], [f385, f1092])).
fof(f14047, plain, (~ spl10_89 | ~ spl10_114), inference(avatar_split_clause, [], [f14042, f1119, f1014])).
fof(f14042, plain, (~ (e3 = op(e1, e2)) | ~ spl10_114), inference(backward_demodulation, [], [f374, f1121])).
fof(f14031, plain, (~ spl10_5 | ~ spl10_125 | spl10_245), inference(avatar_contradiction_clause, [], [f14030])).
fof(f14030, plain, ($false | (~ spl10_5 | ~ spl10_125 | spl10_245)), inference(subsumption_resolution, [], [f14021, f663])).
fof(f14021, plain, (~ (e4 = op(e4, e4)) | (~ spl10_125 | spl10_245)), inference(backward_demodulation, [], [f1861, f1167])).
fof(f1861, plain, (~ (e4 = op(op(e0, e0), op(e0, e0))) | spl10_245), inference(avatar_component_clause, [], [f1859])).
fof(f14023, plain, (~ spl10_115 | ~ spl10_125), inference(avatar_split_clause, [], [f14011, f1165, f1123])).
fof(f14011, plain, (~ (e4 = op(e0, e2)) | ~ spl10_125), inference(backward_demodulation, [], [f405, f1167])).
fof(f14022, plain, (~ spl10_50 | ~ spl10_125), inference(avatar_split_clause, [], [f14009, f1165, f850])).
fof(f14009, plain, (~ (e4 = op(e3, e0)) | ~ spl10_125), inference(backward_demodulation, [], [f356, f1167])).
fof(f14000, plain, (~ spl10_43 | ~ spl10_37 | spl10_172), inference(avatar_split_clause, [], [f13999, f1443, f796, f821])).
fof(f13999, plain, (~ (e2 = op(e3, e1)) | (~ spl10_37 | spl10_172)), inference(forward_demodulation, [], [f1445, f798])).
fof(f13998, plain, (~ spl10_50 | ~ spl10_31 | spl10_193), inference(avatar_split_clause, [], [f13997, f1552, f771, f850])).
fof(f13997, plain, (~ (e4 = op(e3, e0)) | (~ spl10_31 | spl10_193)), inference(forward_demodulation, [], [f1554, f773])).
fof(f773, plain, ((e0 = op(e3, e3)) | ~ spl10_31), inference(avatar_component_clause, [], [f771])).
fof(f13994, plain, (spl10_99 | ~ spl10_91 | ~ spl10_200), inference(avatar_split_clause, [], [f13993, f1583, f1023, f1056])).
fof(f13993, plain, ((e3 = op(e1, e0)) | (~ spl10_91 | ~ spl10_200)), inference(forward_demodulation, [], [f1584, f1025])).
fof(f13991, plain, (~ spl10_98 | ~ spl10_91 | spl10_207), inference(avatar_split_clause, [], [f13990, f1617, f1023, f1052])).
fof(f13990, plain, (~ (e2 = op(e1, e0)) | (~ spl10_91 | spl10_207)), inference(forward_demodulation, [], [f1619, f1025])).
fof(f13982, plain, (~ spl10_124 | ~ spl10_91 | spl10_238), inference(avatar_split_clause, [], [f13670, f1801, f1023, f1161])).
fof(f1801, plain, (spl10_238 <=> (e3 = op(op(e1, e1), op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl10_238])])).
fof(f13670, plain, (~ (op(e0, e0) = e3) | (~ spl10_91 | spl10_238)), inference(backward_demodulation, [], [f1803, f1025])).
fof(f1803, plain, (~ (e3 = op(op(e1, e1), op(e1, e1))) | spl10_238), inference(avatar_component_clause, [], [f1801])).
fof(f13981, plain, (~ spl10_125 | ~ spl10_91 | spl10_242), inference(avatar_split_clause, [], [f13671, f1841, f1023, f1165])).
fof(f13671, plain, (~ (op(e0, e0) = e4) | (~ spl10_91 | spl10_242)), inference(backward_demodulation, [], [f1843, f1025])).
fof(f13954, plain, (~ spl10_74 | ~ spl10_61 | spl10_188), inference(avatar_split_clause, [], [f13890, f1529, f897, f951])).
fof(f13890, plain, (~ (e3 = op(e2, e0)) | (~ spl10_61 | spl10_188)), inference(backward_demodulation, [], [f1531, f899])).
fof(f1531, plain, (~ (e3 = op(e2, op(e2, e2))) | spl10_188), inference(avatar_component_clause, [], [f1529])).
fof(f13953, plain, (~ spl10_75 | ~ spl10_31 | ~ spl10_192 | spl10_213), inference(avatar_split_clause, [], [f13932, f1651, f1547, f771, f955])).
fof(f1547, plain, (spl10_192 <=> (e2 = op(e3, op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl10_192])])).
fof(f13932, plain, (~ (e4 = op(e2, e0)) | (~ spl10_31 | ~ spl10_192 | spl10_213)), inference(backward_demodulation, [], [f13866, f773])).
fof(f13866, plain, (~ (e4 = op(e2, op(e3, e3))) | (~ spl10_192 | spl10_213)), inference(forward_demodulation, [], [f1653, f1548])).
fof(f1548, plain, ((e2 = op(e3, op(e3, e3))) | ~ spl10_192), inference(avatar_component_clause, [], [f1547])).
fof(f13951, plain, (~ spl10_72 | ~ spl10_31 | spl10_191 | ~ spl10_192), inference(avatar_split_clause, [], [f13931, f1547, f1543, f771, f943])).
fof(f1543, plain, (spl10_191 <=> (e1 = op(op(e3, op(e3, e3)), op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl10_191])])).
fof(f13931, plain, (~ (e1 = op(e2, e0)) | (~ spl10_31 | spl10_191 | ~ spl10_192)), inference(backward_demodulation, [], [f13834, f773])).
fof(f13834, plain, (~ (e1 = op(e2, op(e3, e3))) | (spl10_191 | ~ spl10_192)), inference(backward_demodulation, [], [f1545, f1548])).
fof(f1545, plain, (~ (e1 = op(op(e3, op(e3, e3)), op(e3, e3))) | spl10_191), inference(avatar_component_clause, [], [f1543])).
fof(f13947, plain, (spl10_69 | ~ spl10_37 | ~ spl10_170), inference(avatar_split_clause, [], [f13916, f1433, f796, f930])).
fof(f13916, plain, ((e3 = op(e2, e1)) | (~ spl10_37 | ~ spl10_170)), inference(backward_demodulation, [], [f1435, f798])).
fof(f1435, plain, ((e3 = op(e2, op(e3, e2))) | ~ spl10_170), inference(avatar_component_clause, [], [f1433])).
fof(f13898, plain, (~ spl10_56 | ~ spl10_61), inference(avatar_split_clause, [], [f13882, f897, f876])).
fof(f13882, plain, (~ (e0 = op(e2, e3)) | ~ spl10_61), inference(backward_demodulation, [], [f431, f899])).
fof(f13873, plain, (~ spl10_56 | ~ spl10_74 | spl10_183), inference(avatar_split_clause, [], [f13872, f1502, f951, f876])).
fof(f13872, plain, (~ (e0 = op(e2, e3)) | (~ spl10_74 | spl10_183)), inference(forward_demodulation, [], [f1504, f953])).
fof(f13840, plain, (spl10_65 | ~ spl10_13 | ~ spl10_171), inference(avatar_split_clause, [], [f13805, f1437, f695, f913])).
fof(f13805, plain, ((e4 = op(e2, e2)) | (~ spl10_13 | ~ spl10_171)), inference(backward_demodulation, [], [f1439, f697])).
fof(f13837, plain, (~ spl10_37 | ~ spl10_43 | spl10_178), inference(avatar_split_clause, [], [f13597, f1475, f821, f796])).
fof(f1475, plain, (spl10_178 <=> (e1 = op(e3, op(e3, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl10_178])])).
fof(f13597, plain, (~ (e1 = op(e3, e2)) | (~ spl10_43 | spl10_178)), inference(forward_demodulation, [], [f1477, f823])).
fof(f1477, plain, (~ (e1 = op(e3, op(e3, e1))) | spl10_178), inference(avatar_component_clause, [], [f1475])).
fof(f13833, plain, (~ spl10_34 | ~ spl10_9), inference(avatar_split_clause, [], [f13832, f678, f783])).
fof(f783, plain, (spl10_34 <=> (e3 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl10_34])])).
fof(f13832, plain, (~ (e3 = op(e3, e3)) | ~ spl10_9), inference(forward_demodulation, [], [f393, f680])).
fof(f680, plain, ((e3 = op(e4, e3)) | ~ spl10_9), inference(avatar_component_clause, [], [f678])).
fof(f13831, plain, (~ spl10_34 | ~ spl10_29), inference(avatar_split_clause, [], [f13830, f762, f783])).
fof(f13830, plain, (~ (e3 = op(e3, e3)) | ~ spl10_29), inference(forward_demodulation, [], [f443, f764])).
fof(f13723, plain, (~ spl10_55 | ~ spl10_70), inference(avatar_split_clause, [], [f13720, f934, f871])).
fof(f13720, plain, (~ (e4 = op(e2, e4)) | ~ spl10_70), inference(backward_demodulation, [], [f430, f936])).
fof(f13716, plain, (spl10_56 | ~ spl10_74 | ~ spl10_183), inference(avatar_split_clause, [], [f13713, f1502, f951, f876])).
fof(f13713, plain, ((e0 = op(e2, e3)) | (~ spl10_74 | ~ spl10_183)), inference(backward_demodulation, [], [f1503, f953])).
fof(f1503, plain, ((e0 = op(e2, op(e2, e0))) | ~ spl10_183), inference(avatar_component_clause, [], [f1502])).
fof(f13694, plain, (~ spl10_79 | ~ spl10_85 | spl10_139), inference(avatar_split_clause, [], [f13693, f1278, f997, f972])).
fof(f13693, plain, (~ (e3 = op(e1, e4)) | (~ spl10_85 | spl10_139)), inference(backward_demodulation, [], [f1280, f999])).
fof(f13675, plain, (~ spl10_81 | ~ spl10_91), inference(avatar_split_clause, [], [f13660, f1023, f981])).
fof(f13660, plain, (~ (e0 = op(e1, e3)) | ~ spl10_91), inference(backward_demodulation, [], [f419, f1025])).
fof(f13674, plain, (~ spl10_16 | ~ spl10_91), inference(avatar_split_clause, [], [f13658, f1023, f708])).
fof(f13658, plain, (~ (e0 = op(e4, e1)) | ~ spl10_91), inference(backward_demodulation, [], [f370, f1025])).
fof(f370, plain, ~ (op(e1, e1) = op(e4, e1)), inference(cnf_transformation, [], [f4])).
fof(f13643, plain, (spl10_36 | ~ spl10_108 | ~ spl10_141), inference(avatar_split_clause, [], [f13639, f1287, f1094, f792])).
fof(f13639, plain, ((e0 = op(e3, e2)) | (~ spl10_108 | ~ spl10_141)), inference(backward_demodulation, [], [f1289, f1096])).
fof(f1289, plain, ((e0 = op(e3, op(e0, e3))) | ~ spl10_141), inference(avatar_component_clause, [], [f1287])).
fof(f13641, plain, (~ spl10_8 | ~ spl10_108), inference(avatar_split_clause, [], [f13635, f1094, f674])).
fof(f13635, plain, (~ (e2 = op(e4, e3)) | ~ spl10_108), inference(backward_demodulation, [], [f387, f1096])).
fof(f13632, plain, (~ spl10_103 | ~ spl10_115 | spl10_148), inference(avatar_split_clause, [], [f13629, f1320, f1123, f1073])).
fof(f13629, plain, (~ (e2 = op(e0, e4)) | (~ spl10_115 | spl10_148)), inference(backward_demodulation, [], [f1322, f1125])).
fof(f13631, plain, (~ spl10_110 | ~ spl10_115), inference(avatar_split_clause, [], [f13626, f1123, f1102])).
fof(f13626, plain, (~ (e4 = op(e0, e3)) | ~ spl10_115), inference(backward_demodulation, [], [f411, f1125])).
fof(f13630, plain, (~ spl10_15 | ~ spl10_115), inference(avatar_split_clause, [], [f13625, f1123, f703])).
fof(f13625, plain, (~ (e4 = op(e4, e2)) | ~ spl10_115), inference(backward_demodulation, [], [f377, f1125])).
fof(f13601, plain, (~ spl10_81 | ~ spl10_119 | spl10_152), inference(avatar_split_clause, [], [f13600, f1339, f1140, f981])).
fof(f13600, plain, (~ (e0 = op(e1, e3)) | (~ spl10_119 | spl10_152)), inference(forward_demodulation, [], [f1340, f1142])).
fof(f13599, plain, (spl10_89 | ~ spl10_43 | ~ spl10_175), inference(avatar_split_clause, [], [f13598, f1461, f821, f1014])).
fof(f13598, plain, ((e3 = op(e1, e2)) | (~ spl10_43 | ~ spl10_175)), inference(forward_demodulation, [], [f1463, f823])).
fof(f13596, plain, (~ spl10_50 | spl10_104 | ~ spl10_181), inference(avatar_contradiction_clause, [], [f13595])).
fof(f13595, plain, ($false | (~ spl10_50 | spl10_104 | ~ spl10_181)), inference(subsumption_resolution, [], [f13594, f1078])).
fof(f1078, plain, (~ (e3 = op(e0, e4)) | spl10_104), inference(avatar_component_clause, [], [f1077])).
fof(f13594, plain, ((e3 = op(e0, e4)) | (~ spl10_50 | ~ spl10_181)), inference(forward_demodulation, [], [f1495, f852])).
fof(f13593, plain, (~ spl10_26 | ~ spl10_50 | spl10_184), inference(avatar_split_clause, [], [f13592, f1507, f850, f750])).
fof(f13592, plain, (~ (e0 = op(e3, e4)) | (~ spl10_50 | spl10_184)), inference(forward_demodulation, [], [f1509, f852])).
fof(f13578, plain, (~ spl10_95 | ~ spl10_32 | spl10_244), inference(avatar_split_clause, [], [f13577, f1853, f775, f1039])).
fof(f1853, plain, (spl10_244 <=> (e4 = op(op(e3, e3), op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl10_244])])).
fof(f13577, plain, (~ (e4 = op(e1, e1)) | (~ spl10_32 | spl10_244)), inference(forward_demodulation, [], [f1855, f777])).
fof(f1855, plain, (~ (e4 = op(op(e3, e3), op(e3, e3))) | spl10_244), inference(avatar_component_clause, [], [f1853])).
fof(f13558, plain, (~ spl10_91 | ~ spl10_32 | spl10_190), inference(avatar_split_clause, [], [f13530, f1539, f775, f1023])).
fof(f13530, plain, (~ (e0 = op(e1, e1)) | (~ spl10_32 | spl10_190)), inference(backward_demodulation, [], [f1541, f777])).
fof(f13550, plain, (~ spl10_40 | ~ spl10_50), inference(avatar_split_clause, [], [f13549, f850, f808])).
fof(f13549, plain, (~ (e4 = op(e3, e2)) | ~ spl10_50), inference(forward_demodulation, [], [f435, f852])).
fof(f13548, plain, (~ spl10_27 | ~ spl10_32), inference(avatar_split_clause, [], [f13547, f775, f754])).
fof(f13547, plain, (~ (e1 = op(e3, e4)) | ~ spl10_32), inference(forward_demodulation, [], [f443, f777])).
fof(f13546, plain, (~ spl10_30 | ~ spl10_50), inference(avatar_split_clause, [], [f13545, f850, f766])).
fof(f13545, plain, (~ (e4 = op(e3, e4)) | ~ spl10_50), inference(forward_demodulation, [], [f437, f852])).
fof(f13544, plain, (~ spl10_25 | ~ spl10_50), inference(avatar_split_clause, [], [f13543, f850, f745])).
fof(f13543, plain, (~ (e4 = op(e4, e0)) | ~ spl10_50), inference(forward_demodulation, [], [f363, f852])).
fof(f363, plain, ~ (op(e3, e0) = op(e4, e0)), inference(cnf_transformation, [], [f4])).
fof(f13515, plain, (spl10_118 | ~ spl10_122 | ~ spl10_227), inference(avatar_contradiction_clause, [], [f13514])).
fof(f13514, plain, ($false | (spl10_118 | ~ spl10_122 | ~ spl10_227)), inference(subsumption_resolution, [], [f13506, f1137])).
fof(f1137, plain, (~ (e2 = op(e0, e1)) | spl10_118), inference(avatar_component_clause, [], [f1136])).
fof(f13506, plain, ((e2 = op(e0, e1)) | (~ spl10_122 | ~ spl10_227)), inference(backward_demodulation, [], [f1722, f1155])).
fof(f13497, plain, (~ spl10_107 | ~ spl10_119 | spl10_153), inference(avatar_split_clause, [], [f13496, f1344, f1140, f1090])).
fof(f1344, plain, (spl10_153 <=> (e1 = op(e0, op(e0, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl10_153])])).
fof(f13496, plain, (~ (e1 = op(e0, e3)) | (~ spl10_119 | spl10_153)), inference(forward_demodulation, [], [f1346, f1142])).
fof(f1346, plain, (~ (e1 = op(e0, op(e0, e1))) | spl10_153), inference(avatar_component_clause, [], [f1344])).
fof(f13455, plain, (~ spl10_10 | ~ spl10_15), inference(avatar_split_clause, [], [f13453, f703, f682])).
fof(f13453, plain, (~ (e4 = op(e4, e3)) | ~ spl10_15), inference(backward_demodulation, [], [f451, f705])).
fof(f13452, plain, (~ spl10_2 | ~ spl10_16 | ~ spl10_119 | spl10_210), inference(avatar_contradiction_clause, [], [f13451])).
fof(f13451, plain, ($false | (~ spl10_2 | ~ spl10_16 | ~ spl10_119 | spl10_210)), inference(subsumption_resolution, [], [f13446, f1142])).
fof(f13446, plain, (~ (e3 = op(e0, e1)) | (~ spl10_2 | ~ spl10_16 | spl10_210)), inference(backward_demodulation, [], [f13134, f710])).
fof(f13134, plain, (~ (e3 = op(op(e4, e1), e1)) | (~ spl10_2 | spl10_210)), inference(backward_demodulation, [], [f1635, f651])).
fof(f13447, plain, (~ spl10_11 | ~ spl10_16), inference(avatar_split_clause, [], [f13437, f708, f687])).
fof(f687, plain, (spl10_11 <=> (e0 = op(e4, e2))), introduced(avatar_definition, [new_symbols(naming, [spl10_11])])).
fof(f13437, plain, (~ (e0 = op(e4, e2)) | ~ spl10_16), inference(backward_demodulation, [], [f448, f710])).
fof(f13426, plain, (~ spl10_2 | ~ spl10_35 | spl10_216), inference(avatar_contradiction_clause, [], [f13425])).
fof(f13425, plain, ($false | (~ spl10_2 | ~ spl10_35 | spl10_216)), inference(subsumption_resolution, [], [f13416, f651])).
fof(f13416, plain, (~ (e1 = op(e4, e4)) | (~ spl10_35 | spl10_216)), inference(backward_demodulation, [], [f1669, f789])).
fof(f1669, plain, (~ (e1 = op(op(e3, e3), op(e3, e3))) | spl10_216), inference(avatar_component_clause, [], [f1667])).
fof(f1667, plain, (spl10_216 <=> (e1 = op(op(e3, e3), op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl10_216])])).
fof(f13423, plain, (~ spl10_10 | ~ spl10_35), inference(avatar_split_clause, [], [f13413, f787, f682])).
fof(f13413, plain, (~ (e4 = op(e4, e3)) | ~ spl10_35), inference(backward_demodulation, [], [f393, f789])).
fof(f13412, plain, (~ spl10_34 | ~ spl10_39), inference(avatar_split_clause, [], [f13407, f804, f783])).
fof(f13407, plain, (~ (e3 = op(e3, e3)) | ~ spl10_39), inference(backward_demodulation, [], [f441, f806])).
fof(f13404, plain, (~ spl10_38 | ~ spl10_43), inference(avatar_split_clause, [], [f13397, f821, f800])).
fof(f13397, plain, (~ (e2 = op(e3, e2)) | ~ spl10_43), inference(backward_demodulation, [], [f438, f823])).
fof(f13395, plain, (~ spl10_42 | ~ spl10_47), inference(avatar_split_clause, [], [f13386, f838, f817])).
fof(f817, plain, (spl10_42 <=> (e1 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl10_42])])).
fof(f13386, plain, (~ (e1 = op(e3, e1)) | ~ spl10_47), inference(backward_demodulation, [], [f434, f840])).
fof(f13384, plain, (~ spl10_30 | ~ spl10_55), inference(avatar_split_clause, [], [f13381, f871, f766])).
fof(f13381, plain, (~ (e4 = op(e3, e4)) | ~ spl10_55), inference(backward_demodulation, [], [f401, f873])).
fof(f401, plain, ~ (op(e2, e4) = op(e3, e4)), inference(cnf_transformation, [], [f4])).
fof(f13379, plain, (~ spl10_34 | ~ spl10_59), inference(avatar_split_clause, [], [f13373, f888, f783])).
fof(f13373, plain, (~ (e3 = op(e3, e3)) | ~ spl10_59), inference(backward_demodulation, [], [f391, f890])).
fof(f13369, plain, (~ spl10_58 | ~ spl10_63), inference(avatar_split_clause, [], [f13358, f905, f884])).
fof(f13358, plain, (~ (e2 = op(e2, e3)) | ~ spl10_63), inference(backward_demodulation, [], [f431, f907])).
fof(f13354, plain, (~ spl10_42 | ~ spl10_67), inference(avatar_split_clause, [], [f13344, f922, f817])).
fof(f13344, plain, (~ (e1 = op(e3, e1)) | ~ spl10_67), inference(backward_demodulation, [], [f371, f924])).
fof(f371, plain, ~ (op(e2, e1) = op(e3, e1)), inference(cnf_transformation, [], [f4])).
fof(f13342, plain, (~ spl10_46 | ~ spl10_71), inference(avatar_split_clause, [], [f13331, f939, f834])).
fof(f834, plain, (spl10_46 <=> (e0 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl10_46])])).
fof(f13331, plain, (~ (e0 = op(e3, e0)) | ~ spl10_71), inference(backward_demodulation, [], [f361, f941])).
fof(f13308, plain, (~ spl10_82 | ~ spl10_87), inference(avatar_split_clause, [], [f13300, f1006, f985])).
fof(f985, plain, (spl10_82 <=> (e1 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl10_82])])).
fof(f13300, plain, (~ (e1 = op(e1, e3)) | ~ spl10_87), inference(backward_demodulation, [], [f421, f1008])).
fof(f13276, plain, (~ spl10_107 | ~ spl10_119 | spl10_142), inference(avatar_contradiction_clause, [], [f13275])).
fof(f13275, plain, ($false | (~ spl10_107 | ~ spl10_119 | spl10_142)), inference(subsumption_resolution, [], [f13273, f1142])).
fof(f13273, plain, (~ (e3 = op(e0, e1)) | (~ spl10_107 | spl10_142)), inference(backward_demodulation, [], [f1294, f1092])).
fof(f13274, plain, (~ spl10_82 | ~ spl10_107), inference(avatar_split_clause, [], [f13265, f1090, f985])).
fof(f13265, plain, (~ (e1 = op(e1, e3)) | ~ spl10_107), inference(backward_demodulation, [], [f384, f1092])).
fof(f13263, plain, (~ spl10_106 | ~ spl10_111), inference(avatar_split_clause, [], [f13255, f1107, f1086])).
fof(f1086, plain, (spl10_106 <=> (e0 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl10_106])])).
fof(f13255, plain, (~ (e0 = op(e0, e3)) | ~ spl10_111), inference(backward_demodulation, [], [f411, f1109])).
fof(f13262, plain, (~ spl10_11 | ~ spl10_111), inference(avatar_split_clause, [], [f13254, f1107, f687])).
fof(f13254, plain, (~ (e0 = op(e4, e2)) | ~ spl10_111), inference(backward_demodulation, [], [f377, f1109])).
fof(f13250, plain, (spl10_107 | ~ spl10_119 | ~ spl10_153), inference(avatar_split_clause, [], [f13246, f1344, f1140, f1090])).
fof(f13246, plain, ((e1 = op(e0, e3)) | (~ spl10_119 | ~ spl10_153)), inference(backward_demodulation, [], [f1345, f1142])).
fof(f1345, plain, ((e1 = op(e0, op(e0, e1))) | ~ spl10_153), inference(avatar_component_clause, [], [f1344])).
fof(f13248, plain, (~ spl10_114 | ~ spl10_119), inference(avatar_split_clause, [], [f13242, f1140, f1119])).
fof(f13242, plain, (~ (e3 = op(e0, e2)) | ~ spl10_119), inference(backward_demodulation, [], [f408, f1142])).
fof(f13233, plain, (~ spl10_2 | ~ spl10_125 | spl10_220), inference(avatar_contradiction_clause, [], [f13232])).
fof(f13232, plain, ($false | (~ spl10_2 | ~ spl10_125 | spl10_220)), inference(subsumption_resolution, [], [f13223, f651])).
fof(f13223, plain, (~ (e1 = op(e4, e4)) | (~ spl10_125 | spl10_220)), inference(backward_demodulation, [], [f1689, f1167])).
fof(f1689, plain, (~ (e1 = op(op(e0, e0), op(e0, e0))) | spl10_220), inference(avatar_component_clause, [], [f1687])).
fof(f1687, plain, (spl10_220 <=> (e1 = op(op(e0, e0), op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl10_220])])).
fof(f13231, plain, (~ spl10_120 | ~ spl10_125), inference(avatar_split_clause, [], [f13220, f1165, f1144])).
fof(f13220, plain, (~ (e4 = op(e0, e1)) | ~ spl10_125), inference(backward_demodulation, [], [f404, f1167])).
fof(f13230, plain, (~ spl10_75 | ~ spl10_125), inference(avatar_split_clause, [], [f13218, f1165, f955])).
fof(f13218, plain, (~ (e4 = op(e2, e0)) | ~ spl10_125), inference(backward_demodulation, [], [f355, f1167])).
fof(f13167, plain, (~ spl10_110 | ~ spl10_24 | spl10_182), inference(avatar_split_clause, [], [f13117, f1497, f741, f1102])).
fof(f13117, plain, (~ (e4 = op(e0, e3)) | (~ spl10_24 | spl10_182)), inference(backward_demodulation, [], [f1498, f743])).
fof(f13166, plain, (~ spl10_94 | ~ spl10_2 | spl10_240), inference(avatar_split_clause, [], [f13138, f1813, f649, f1035])).
fof(f1813, plain, (spl10_240 <=> (e3 = op(op(e4, e4), op(e4, e4)))), introduced(avatar_definition, [new_symbols(naming, [spl10_240])])).
fof(f13138, plain, (~ (e3 = op(e1, e1)) | (~ spl10_2 | spl10_240)), inference(backward_demodulation, [], [f1815, f651])).
fof(f1815, plain, (~ (e3 = op(op(e4, e4), op(e4, e4))) | spl10_240), inference(avatar_component_clause, [], [f1813])).
fof(f13164, plain, (~ spl10_78 | ~ spl10_103), inference(avatar_split_clause, [], [f13163, f1073, f968])).
fof(f13163, plain, (~ (e2 = op(e1, e4)) | ~ spl10_103), inference(forward_demodulation, [], [f394, f1075])).
fof(f1075, plain, ((e2 = op(e0, e4)) | ~ spl10_103), inference(avatar_component_clause, [], [f1073])).
fof(f13155, plain, (~ spl10_18 | ~ spl10_2 | spl10_196), inference(avatar_split_clause, [], [f13131, f1565, f649, f716])).
fof(f1565, plain, (spl10_196 <=> (e2 = op(e4, op(e4, e4)))), introduced(avatar_definition, [new_symbols(naming, [spl10_196])])).
fof(f13131, plain, (~ (e2 = op(e4, e1)) | (~ spl10_2 | spl10_196)), inference(backward_demodulation, [], [f1567, f651])).
fof(f1567, plain, (~ (e2 = op(e4, op(e4, e4))) | spl10_196), inference(avatar_component_clause, [], [f1565])).
fof(f13153, plain, (~ spl10_14 | ~ spl10_24), inference(avatar_split_clause, [], [f13152, f741, f699])).
fof(f13152, plain, (~ (e3 = op(e4, e2)) | ~ spl10_24), inference(forward_demodulation, [], [f445, f743])).
fof(f13147, plain, (~ spl10_2 | ~ spl10_18 | ~ spl10_66 | spl10_219), inference(avatar_contradiction_clause, [], [f13146])).
fof(f13146, plain, ($false | (~ spl10_2 | ~ spl10_18 | ~ spl10_66 | spl10_219)), inference(subsumption_resolution, [], [f13145, f920])).
fof(f13145, plain, (~ (e0 = op(e2, e1)) | (~ spl10_2 | ~ spl10_18 | spl10_219)), inference(forward_demodulation, [], [f13136, f718])).
fof(f13136, plain, (~ (e0 = op(op(e4, e1), e1)) | (~ spl10_2 | spl10_219)), inference(backward_demodulation, [], [f1683, f651])).
fof(f1683, plain, (~ (e0 = op(op(e4, op(e4, e4)), op(e4, e4))) | spl10_219), inference(avatar_component_clause, [], [f1681])).
fof(f1681, plain, (spl10_219 <=> (e0 = op(op(e4, op(e4, e4)), op(e4, e4)))), introduced(avatar_definition, [new_symbols(naming, [spl10_219])])).
fof(f13126, plain, (~ spl10_3 | ~ spl10_18), inference(avatar_split_clause, [], [f13121, f716, f653])).
fof(f13121, plain, (~ (e2 = op(e4, e4)) | ~ spl10_18), inference(backward_demodulation, [], [f450, f718])).
fof(f450, plain, ~ (op(e4, e1) = op(e4, e4)), inference(cnf_transformation, [], [f4])).
fof(f13105, plain, (~ spl10_62 | ~ spl10_94 | spl10_239), inference(avatar_contradiction_clause, [], [f13104])).
fof(f13104, plain, ($false | (~ spl10_62 | ~ spl10_94 | spl10_239)), inference(subsumption_resolution, [], [f13100, f1037])).
fof(f13100, plain, (~ (e3 = op(e1, e1)) | (~ spl10_62 | spl10_239)), inference(backward_demodulation, [], [f1809, f903])).
fof(f1809, plain, (~ (e3 = op(op(e2, e2), op(e2, e2))) | spl10_239), inference(avatar_component_clause, [], [f1807])).
fof(f1807, plain, (spl10_239 <=> (e3 = op(op(e2, e2), op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl10_239])])).
fof(f13103, plain, (~ spl10_52 | ~ spl10_62), inference(avatar_split_clause, [], [f13093, f901, f859])).
fof(f859, plain, (spl10_52 <=> (e1 = op(e2, e4))), introduced(avatar_definition, [new_symbols(naming, [spl10_52])])).
fof(f13093, plain, (~ (e1 = op(e2, e4)) | ~ spl10_62), inference(backward_demodulation, [], [f432, f903])).
fof(f13075, plain, (~ spl10_19 | ~ spl10_94), inference(avatar_split_clause, [], [f13065, f1035, f720])).
fof(f13065, plain, (~ (e3 = op(e4, e1)) | ~ spl10_94), inference(backward_demodulation, [], [f370, f1037])).
fof(f13063, plain, (~ spl10_3 | ~ spl10_103), inference(avatar_split_clause, [], [f13058, f1073, f653])).
fof(f13058, plain, (~ (e2 = op(e4, e4)) | ~ spl10_103), inference(backward_demodulation, [], [f397, f1075])).
fof(f13047, plain, (~ spl10_19 | ~ spl10_120 | ~ spl10_122 | spl10_226), inference(avatar_split_clause, [], [f13043, f1717, f1153, f1144, f720])).
fof(f13043, plain, (~ (e3 = op(e4, e1)) | (~ spl10_120 | ~ spl10_122 | spl10_226)), inference(backward_demodulation, [], [f13028, f1146])).
fof(f13028, plain, (~ (e3 = op(op(e0, e1), e1)) | (~ spl10_122 | spl10_226)), inference(backward_demodulation, [], [f1719, f1155])).
fof(f13046, plain, (~ spl10_18 | ~ spl10_120 | ~ spl10_122 | spl10_221), inference(avatar_split_clause, [], [f13042, f1691, f1153, f1144, f716])).
fof(f13042, plain, (~ (e2 = op(e4, e1)) | (~ spl10_120 | ~ spl10_122 | spl10_221)), inference(backward_demodulation, [], [f13026, f1146])).
fof(f13026, plain, (~ (e2 = op(op(e0, e1), e1)) | (~ spl10_122 | spl10_221)), inference(backward_demodulation, [], [f1693, f1155])).
fof(f13045, plain, (~ spl10_95 | ~ spl10_120), inference(avatar_split_clause, [], [f13037, f1144, f1039])).
fof(f13037, plain, (~ (e4 = op(e1, e1)) | ~ spl10_120), inference(backward_demodulation, [], [f364, f1146])).
fof(f13035, plain, (~ spl10_120 | ~ spl10_122 | spl10_223), inference(avatar_split_clause, [], [f13027, f1700, f1153, f1144])).
fof(f1700, plain, (spl10_223 <=> (e4 = op(e0, op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl10_223])])).
fof(f13027, plain, (~ (e4 = op(e0, e1)) | (~ spl10_122 | spl10_223)), inference(backward_demodulation, [], [f1702, f1155])).
fof(f1702, plain, (~ (e4 = op(e0, op(e0, e0))) | spl10_223), inference(avatar_component_clause, [], [f1700])).
fof(f13033, plain, (~ spl10_22 | ~ spl10_122), inference(avatar_split_clause, [], [f13018, f1153, f733])).
fof(f733, plain, (spl10_22 <=> (e1 = op(e4, e0))), introduced(avatar_definition, [new_symbols(naming, [spl10_22])])).
fof(f13018, plain, (~ (e1 = op(e4, e0)) | ~ spl10_122), inference(backward_demodulation, [], [f357, f1155])).
fof(f13009, plain, (spl10_75 | ~ spl10_11 | ~ spl10_171), inference(avatar_split_clause, [], [f13008, f1437, f687, f955])).
fof(f13008, plain, ((e4 = op(e2, e0)) | (~ spl10_11 | ~ spl10_171)), inference(forward_demodulation, [], [f1439, f689])).
fof(f689, plain, ((e0 = op(e4, e2)) | ~ spl10_11), inference(avatar_component_clause, [], [f687])).
fof(f12983, plain, (~ spl10_53 | ~ spl10_58), inference(avatar_split_clause, [], [f12982, f884, f863])).
fof(f12982, plain, (~ (e2 = op(e2, e4)) | ~ spl10_58), inference(forward_demodulation, [], [f433, f886])).
fof(f886, plain, ((e2 = op(e2, e3)) | ~ spl10_58), inference(avatar_component_clause, [], [f884])).
fof(f12977, plain, (~ spl10_25 | ~ spl10_10), inference(avatar_split_clause, [], [f12976, f682, f745])).
fof(f12976, plain, (~ (e4 = op(e4, e0)) | ~ spl10_10), inference(forward_demodulation, [], [f446, f684])).
fof(f684, plain, ((e4 = op(e4, e3)) | ~ spl10_10), inference(avatar_component_clause, [], [f682])).
fof(f12975, plain, (~ spl10_23 | ~ spl10_98), inference(avatar_split_clause, [], [f12974, f1052, f737])).
fof(f12974, plain, (~ (e2 = op(e4, e0)) | ~ spl10_98), inference(forward_demodulation, [], [f360, f1054])).
fof(f12959, plain, (~ spl10_1 | ~ spl10_11), inference(avatar_split_clause, [], [f12952, f687, f645])).
fof(f12952, plain, (~ (e0 = op(e4, e4)) | ~ spl10_11), inference(backward_demodulation, [], [f452, f689])).
fof(f12937, plain, (~ spl10_61 | ~ spl10_66), inference(avatar_split_clause, [], [f12930, f918, f897])).
fof(f12930, plain, (~ (e0 = op(e2, e2)) | ~ spl10_66), inference(backward_demodulation, [], [f428, f920])).
fof(f12927, plain, (~ spl10_69 | ~ spl10_74), inference(avatar_split_clause, [], [f12922, f951, f930])).
fof(f12922, plain, (~ (e3 = op(e2, e1)) | ~ spl10_74), inference(backward_demodulation, [], [f424, f953])).
fof(f12903, plain, (spl10_76 | ~ spl10_95 | ~ spl10_237), inference(avatar_split_clause, [], [f12896, f1787, f1039, f960])).
fof(f12896, plain, ((e0 = op(e1, e4)) | (~ spl10_95 | ~ spl10_237)), inference(backward_demodulation, [], [f1788, f1041])).
fof(f12902, plain, (~ spl10_3 | ~ spl10_95 | spl10_230), inference(avatar_split_clause, [], [f12895, f1745, f1039, f653])).
fof(f12895, plain, (~ (e2 = op(e4, e4)) | (~ spl10_95 | spl10_230)), inference(backward_demodulation, [], [f1747, f1041])).
fof(f12901, plain, (~ spl10_1 | ~ spl10_95 | spl10_198), inference(avatar_split_clause, [], [f12894, f1575, f1039, f645])).
fof(f12894, plain, (~ (e0 = op(e4, e4)) | (~ spl10_95 | spl10_198)), inference(backward_demodulation, [], [f1577, f1041])).
fof(f1577, plain, (~ (e0 = op(op(e1, e1), op(e1, e1))) | spl10_198), inference(avatar_component_clause, [], [f1575])).
fof(f12882, plain, (~ spl10_103 | ~ spl10_118), inference(avatar_split_clause, [], [f12876, f1136, f1073])).
fof(f12876, plain, (~ (e2 = op(e0, e4)) | ~ spl10_118), inference(backward_demodulation, [], [f410, f1138])).
fof(f12881, plain, (~ spl10_18 | ~ spl10_118), inference(avatar_split_clause, [], [f12875, f1136, f716])).
fof(f12875, plain, (~ (e2 = op(e4, e1)) | ~ spl10_118), inference(backward_demodulation, [], [f367, f1138])).
fof(f12872, plain, (~ spl10_3 | ~ spl10_125 | spl10_234), inference(avatar_split_clause, [], [f12863, f1767, f1165, f653])).
fof(f12863, plain, (~ (e2 = op(e4, e4)) | (~ spl10_125 | spl10_234)), inference(backward_demodulation, [], [f1769, f1167])).
fof(f12871, plain, (spl10_104 | ~ spl10_125 | ~ spl10_222), inference(avatar_split_clause, [], [f12862, f1695, f1165, f1077])).
fof(f12862, plain, ((e3 = op(e0, e4)) | (~ spl10_125 | ~ spl10_222)), inference(backward_demodulation, [], [f1696, f1167])).
fof(f12855, plain, (~ spl10_118 | ~ spl10_112 | spl10_148), inference(avatar_split_clause, [], [f12854, f1320, f1111, f1136])).
fof(f12854, plain, (~ (e2 = op(e0, e1)) | (~ spl10_112 | spl10_148)), inference(forward_demodulation, [], [f1322, f1113])).
fof(f12819, plain, (~ spl10_94 | ~ spl10_82 | spl10_139), inference(avatar_split_clause, [], [f12669, f1278, f985, f1035])).
fof(f12669, plain, (~ (e3 = op(e1, e1)) | (~ spl10_82 | spl10_139)), inference(backward_demodulation, [], [f1280, f987])).
fof(f987, plain, ((e1 = op(e1, e3)) | ~ spl10_82), inference(avatar_component_clause, [], [f985])).
fof(f12808, plain, (~ spl10_64 | ~ spl10_58 | spl10_167), inference(avatar_split_clause, [], [f12719, f1414, f884, f909])).
fof(f12719, plain, (~ (e3 = op(e2, e2)) | (~ spl10_58 | spl10_167)), inference(backward_demodulation, [], [f1416, f886])).
fof(f12791, plain, (~ spl10_4 | ~ spl10_10 | spl10_168), inference(avatar_split_clause, [], [f12788, f1421, f682, f657])).
fof(f1421, plain, (spl10_168 <=> (e3 = op(e4, op(e4, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl10_168])])).
fof(f12788, plain, (~ (e3 = op(e4, e4)) | (~ spl10_10 | spl10_168)), inference(backward_demodulation, [], [f1423, f684])).
fof(f1423, plain, (~ (e3 = op(e4, op(e4, e3))) | spl10_168), inference(avatar_component_clause, [], [f1421])).
fof(f12785, plain, (~ spl10_8 | ~ spl10_14 | spl10_173), inference(avatar_split_clause, [], [f12782, f1449, f699, f674])).
fof(f12782, plain, (~ (e2 = op(e4, e3)) | (~ spl10_14 | spl10_173)), inference(backward_demodulation, [], [f1451, f701])).
fof(f1451, plain, (~ (e2 = op(e4, op(e4, e2))) | spl10_173), inference(avatar_component_clause, [], [f1449])).
fof(f12783, plain, (~ spl10_4 | ~ spl10_14), inference(avatar_split_clause, [], [f12779, f699, f657])).
fof(f12779, plain, (~ (e3 = op(e4, e4)) | ~ spl10_14), inference(backward_demodulation, [], [f452, f701])).
fof(f12777, plain, (~ spl10_8 | ~ spl10_18), inference(avatar_split_clause, [], [f12771, f716, f674])).
fof(f12771, plain, (~ (e2 = op(e4, e3)) | ~ spl10_18), inference(backward_demodulation, [], [f449, f718])).
fof(f12748, plain, (~ spl10_42 | ~ spl10_91 | ~ spl10_175), inference(avatar_contradiction_clause, [], [f12747])).
fof(f12747, plain, ($false | (~ spl10_42 | ~ spl10_91 | ~ spl10_175)), inference(subsumption_resolution, [], [f12746, f456])).
fof(f456, plain, ~ (e0 = e3), inference(cnf_transformation, [], [f5])).
fof(f12746, plain, ((e0 = e3) | (~ spl10_42 | ~ spl10_91 | ~ spl10_175)), inference(forward_demodulation, [], [f12743, f1025])).
fof(f12743, plain, ((e3 = op(e1, e1)) | (~ spl10_42 | ~ spl10_175)), inference(backward_demodulation, [], [f1463, f819])).
fof(f819, plain, ((e1 = op(e3, e1)) | ~ spl10_42), inference(avatar_component_clause, [], [f817])).
fof(f12745, plain, (~ spl10_37 | ~ spl10_42), inference(avatar_split_clause, [], [f12737, f817, f796])).
fof(f12737, plain, (~ (e1 = op(e3, e2)) | ~ spl10_42), inference(backward_demodulation, [], [f438, f819])).
fof(f12720, plain, (~ spl10_8 | ~ spl10_58), inference(avatar_split_clause, [], [f12715, f884, f674])).
fof(f12715, plain, (~ (e2 = op(e4, e3)) | ~ spl10_58), inference(backward_demodulation, [], [f392, f886])).
fof(f12710, plain, (~ spl10_61 | ~ spl10_124 | spl10_239), inference(avatar_contradiction_clause, [], [f12709])).
fof(f12709, plain, ($false | (~ spl10_61 | ~ spl10_124 | spl10_239)), inference(subsumption_resolution, [], [f12701, f1163])).
fof(f12701, plain, (~ (op(e0, e0) = e3) | (~ spl10_61 | spl10_239)), inference(backward_demodulation, [], [f1809, f899])).
fof(f12705, plain, (~ spl10_11 | ~ spl10_61), inference(avatar_split_clause, [], [f12692, f897, f687])).
fof(f12692, plain, (~ (e0 = op(e4, e2)) | ~ spl10_61), inference(backward_demodulation, [], [f382, f899])).
fof(f12688, plain, (~ spl10_44 | ~ spl10_69), inference(avatar_split_clause, [], [f12681, f930, f825])).
fof(f825, plain, (spl10_44 <=> (e3 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl10_44])])).
fof(f12681, plain, (~ (e3 = op(e3, e1)) | ~ spl10_69), inference(backward_demodulation, [], [f371, f932])).
fof(f12677, plain, (~ spl10_4 | ~ spl10_79), inference(avatar_split_clause, [], [f12673, f972, f657])).
fof(f12673, plain, (~ (e3 = op(e4, e4)) | ~ spl10_79), inference(backward_demodulation, [], [f400, f974])).
fof(f12651, plain, (~ spl10_75 | ~ spl10_91 | ~ spl10_98 | spl10_211), inference(avatar_contradiction_clause, [], [f12650])).
fof(f12650, plain, ($false | (~ spl10_75 | ~ spl10_91 | ~ spl10_98 | spl10_211)), inference(subsumption_resolution, [], [f12649, f957])).
fof(f12649, plain, (~ (e4 = op(e2, e0)) | (~ spl10_91 | ~ spl10_98 | spl10_211)), inference(forward_demodulation, [], [f12639, f1054])).
fof(f12639, plain, (~ (e4 = op(op(e1, e0), e0)) | (~ spl10_91 | spl10_211)), inference(backward_demodulation, [], [f1641, f1025])).
fof(f12623, plain, (~ spl10_88 | ~ spl10_98), inference(avatar_split_clause, [], [f12615, f1052, f1010])).
fof(f1010, plain, (spl10_88 <=> (e2 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl10_88])])).
fof(f12615, plain, (~ (e2 = op(e1, e2)) | ~ spl10_98), inference(backward_demodulation, [], [f415, f1054])).
fof(f12622, plain, (~ spl10_48 | ~ spl10_98), inference(avatar_split_clause, [], [f12613, f1052, f842])).
fof(f12613, plain, (~ (e2 = op(e3, e0)) | ~ spl10_98), inference(backward_demodulation, [], [f359, f1054])).
fof(f12594, plain, (~ spl10_107 | ~ spl10_112), inference(avatar_split_clause, [], [f12587, f1111, f1090])).
fof(f12587, plain, (~ (e1 = op(e0, e3)) | ~ spl10_112), inference(backward_demodulation, [], [f411, f1113])).
fof(f12593, plain, (~ spl10_37 | ~ spl10_112), inference(avatar_split_clause, [], [f12585, f1111, f796])).
fof(f12585, plain, (~ (e1 = op(e3, e2)) | ~ spl10_112), inference(backward_demodulation, [], [f376, f1113])).
fof(f12572, plain, (spl10_10 | ~ spl10_129), inference(avatar_split_clause, [], [f12561, f1182, f682])).
fof(f1182, plain, (spl10_129 <=> (e3 = unit)), introduced(avatar_definition, [new_symbols(naming, [spl10_129])])).
fof(f12561, plain, ((e4 = op(e4, e3)) | ~ spl10_129), inference(backward_demodulation, [], [f302, f1184])).
fof(f1184, plain, ((e3 = unit) | ~ spl10_129), inference(avatar_component_clause, [], [f1182])).
fof(f12571, plain, (spl10_30 | ~ spl10_129), inference(avatar_split_clause, [], [f12560, f1182, f766])).
fof(f12560, plain, ((e4 = op(e3, e4)) | ~ spl10_129), inference(backward_demodulation, [], [f301, f1184])).
fof(f12570, plain, (spl10_34 | ~ spl10_129), inference(avatar_split_clause, [], [f12559, f1182, f783])).
fof(f12559, plain, ((e3 = op(e3, e3)) | ~ spl10_129), inference(backward_demodulation, [], [f300, f1184])).
fof(f12568, plain, (spl10_58 | ~ spl10_129), inference(avatar_split_clause, [], [f12557, f1182, f884])).
fof(f12557, plain, ((e2 = op(e2, e3)) | ~ spl10_129), inference(backward_demodulation, [], [f298, f1184])).
fof(f12567, plain, (spl10_38 | ~ spl10_129), inference(avatar_split_clause, [], [f12556, f1182, f800])).
fof(f12556, plain, ((e2 = op(e3, e2)) | ~ spl10_129), inference(backward_demodulation, [], [f297, f1184])).
fof(f12566, plain, (spl10_82 | ~ spl10_129), inference(avatar_split_clause, [], [f12555, f1182, f985])).
fof(f12555, plain, ((e1 = op(e1, e3)) | ~ spl10_129), inference(backward_demodulation, [], [f296, f1184])).
fof(f12565, plain, (spl10_42 | ~ spl10_129), inference(avatar_split_clause, [], [f12554, f1182, f817])).
fof(f12554, plain, ((e1 = op(e3, e1)) | ~ spl10_129), inference(backward_demodulation, [], [f295, f1184])).
fof(f12564, plain, (spl10_106 | ~ spl10_129), inference(avatar_split_clause, [], [f12553, f1182, f1086])).
fof(f12553, plain, ((e0 = op(e0, e3)) | ~ spl10_129), inference(backward_demodulation, [], [f294, f1184])).
fof(f12563, plain, (spl10_46 | ~ spl10_129), inference(avatar_split_clause, [], [f12552, f1182, f834])).
fof(f12552, plain, ((e0 = op(e3, e0)) | ~ spl10_129), inference(backward_demodulation, [], [f293, f1184])).
fof(f12549, plain, (spl10_18 | ~ spl10_52 | ~ spl10_159), inference(avatar_split_clause, [], [f12548, f1373, f859, f716])).
fof(f12548, plain, ((e2 = op(e4, e1)) | (~ spl10_52 | ~ spl10_159)), inference(forward_demodulation, [], [f1375, f861])).
fof(f861, plain, ((e1 = op(e2, e4)) | ~ spl10_52), inference(avatar_component_clause, [], [f859])).
fof(f12543, plain, (spl10_120 | ~ spl10_22 | ~ spl10_182), inference(avatar_split_clause, [], [f12542, f1497, f733, f1144])).
fof(f12542, plain, ((e4 = op(e0, e1)) | (~ spl10_22 | ~ spl10_182)), inference(forward_demodulation, [], [f1499, f735])).
fof(f735, plain, ((e1 = op(e4, e0)) | ~ spl10_22), inference(avatar_component_clause, [], [f733])).
fof(f12530, plain, (~ spl10_107 | ~ spl10_124 | spl10_236), inference(avatar_split_clause, [], [f12529, f1781, f1161, f1090])).
fof(f12529, plain, (~ (e1 = op(e0, e3)) | (~ spl10_124 | spl10_236)), inference(forward_demodulation, [], [f1783, f1163])).
fof(f1783, plain, (~ (e1 = op(e0, op(e0, e0))) | spl10_236), inference(avatar_component_clause, [], [f1781])).
fof(f12525, plain, (~ spl10_35 | ~ spl10_124 | spl10_245), inference(avatar_split_clause, [], [f12524, f1859, f1161, f787])).
fof(f12524, plain, (~ (e4 = op(e3, e3)) | (~ spl10_124 | spl10_245)), inference(forward_demodulation, [], [f1861, f1163])).
fof(f12515, plain, (~ spl10_109 | ~ spl10_124), inference(avatar_split_clause, [], [f12514, f1161, f1098])).
fof(f12514, plain, (~ (e3 = op(e0, e3)) | ~ spl10_124), inference(forward_demodulation, [], [f406, f1163])).
fof(f12507, plain, (~ spl10_70 | ~ spl10_52 | spl10_162), inference(avatar_split_clause, [], [f12186, f1386, f859, f934])).
fof(f12186, plain, (~ (e4 = op(e2, e1)) | (~ spl10_52 | spl10_162)), inference(backward_demodulation, [], [f1388, f861])).
fof(f12504, plain, (~ spl10_60 | ~ spl10_75), inference(avatar_split_clause, [], [f12503, f955, f892])).
fof(f12503, plain, (~ (e4 = op(e2, e3)) | ~ spl10_75), inference(forward_demodulation, [], [f426, f957])).
fof(f12502, plain, (~ spl10_50 | ~ spl10_75), inference(avatar_split_clause, [], [f12501, f955, f850])).
fof(f12501, plain, (~ (e4 = op(e3, e0)) | ~ spl10_75), inference(forward_demodulation, [], [f361, f957])).
fof(f12500, plain, (~ spl10_49 | ~ spl10_124), inference(avatar_split_clause, [], [f12499, f1161, f846])).
fof(f12499, plain, (~ (e3 = op(e3, e0)) | ~ spl10_124), inference(forward_demodulation, [], [f356, f1163])).
fof(f12495, plain, (~ spl10_35 | ~ spl10_64 | spl10_243), inference(avatar_contradiction_clause, [], [f12494])).
fof(f12494, plain, ($false | (~ spl10_35 | ~ spl10_64 | spl10_243)), inference(subsumption_resolution, [], [f12489, f789])).
fof(f12489, plain, (~ (e4 = op(e3, e3)) | (~ spl10_64 | spl10_243)), inference(backward_demodulation, [], [f1849, f911])).
fof(f12442, plain, (~ spl10_26 | ~ spl10_35 | ~ spl10_103 | spl10_202), inference(avatar_contradiction_clause, [], [f12441])).
fof(f12441, plain, ($false | (~ spl10_26 | ~ spl10_35 | ~ spl10_103 | spl10_202)), inference(subsumption_resolution, [], [f12440, f1075])).
fof(f12440, plain, (~ (e2 = op(e0, e4)) | (~ spl10_26 | ~ spl10_35 | spl10_202)), inference(forward_demodulation, [], [f12439, f752])).
fof(f12439, plain, (~ (e2 = op(op(e3, e4), e4)) | (~ spl10_35 | spl10_202)), inference(forward_demodulation, [], [f1595, f789])).
fof(f12434, plain, (~ spl10_115 | ~ spl10_103 | spl10_136), inference(avatar_split_clause, [], [f12349, f1264, f1073, f1123])).
fof(f12349, plain, (~ (e4 = op(e0, e2)) | (~ spl10_103 | spl10_136)), inference(backward_demodulation, [], [f1266, f1075])).
fof(f12429, plain, (~ spl10_74 | ~ spl10_56 | spl10_167), inference(avatar_split_clause, [], [f12375, f1414, f876, f951])).
fof(f12375, plain, (~ (e3 = op(e2, e0)) | (~ spl10_56 | spl10_167)), inference(backward_demodulation, [], [f1416, f878])).
fof(f12405, plain, (~ spl10_1 | ~ spl10_26), inference(avatar_split_clause, [], [f12398, f750, f645])).
fof(f12398, plain, (~ (e0 = op(e4, e4)) | ~ spl10_26), inference(backward_demodulation, [], [f403, f752])).
fof(f12394, plain, (~ spl10_1 | ~ spl10_35 | spl10_190), inference(avatar_split_clause, [], [f12388, f1539, f787, f645])).
fof(f12388, plain, (~ (e0 = op(e4, e4)) | (~ spl10_35 | spl10_190)), inference(backward_demodulation, [], [f1541, f789])).
fof(f12386, plain, (~ spl10_28 | ~ spl10_48), inference(avatar_split_clause, [], [f12381, f842, f758])).
fof(f12381, plain, (~ (e2 = op(e3, e4)) | ~ spl10_48), inference(backward_demodulation, [], [f437, f844])).
fof(f12376, plain, (~ spl10_31 | ~ spl10_56), inference(avatar_split_clause, [], [f12369, f876, f771])).
fof(f12369, plain, (~ (e0 = op(e3, e3)) | ~ spl10_56), inference(backward_demodulation, [], [f391, f878])).
fof(f12365, plain, (~ spl10_1 | ~ spl10_65 | spl10_186), inference(avatar_split_clause, [], [f12355, f1521, f913, f645])).
fof(f12355, plain, (~ (e0 = op(e4, e4)) | (~ spl10_65 | spl10_186)), inference(backward_demodulation, [], [f1523, f915])).
fof(f1523, plain, (~ (e0 = op(op(e2, e2), op(e2, e2))) | spl10_186), inference(avatar_component_clause, [], [f1521])).
fof(f12351, plain, (spl10_11 | ~ spl10_103 | ~ spl10_135), inference(avatar_split_clause, [], [f12348, f1259, f1073, f687])).
fof(f12348, plain, ((e0 = op(e4, e2)) | (~ spl10_103 | ~ spl10_135)), inference(backward_demodulation, [], [f1261, f1075])).
fof(f12350, plain, (~ spl10_28 | ~ spl10_103), inference(avatar_split_clause, [], [f12344, f1073, f758])).
fof(f12344, plain, (~ (e2 = op(e3, e4)) | ~ spl10_103), inference(backward_demodulation, [], [f396, f1075])).
fof(f12340, plain, (~ spl10_14 | ~ spl10_114), inference(avatar_split_clause, [], [f12335, f1119, f699])).
fof(f12335, plain, (~ (e3 = op(e4, e2)) | ~ spl10_114), inference(backward_demodulation, [], [f377, f1121])).
fof(f12331, plain, (spl10_4 | ~ spl10_125 | ~ spl10_241), inference(avatar_split_clause, [], [f12323, f1819, f1165, f657])).
fof(f12323, plain, ((e3 = op(e4, e4)) | (~ spl10_125 | ~ spl10_241)), inference(backward_demodulation, [], [f1820, f1167])).
fof(f12330, plain, (~ spl10_103 | ~ spl10_125 | spl10_227), inference(avatar_split_clause, [], [f12321, f1721, f1165, f1073])).
fof(f12321, plain, (~ (e2 = op(e0, e4)) | (~ spl10_125 | spl10_227)), inference(backward_demodulation, [], [f1723, f1167])).
fof(f12329, plain, (spl10_105 | ~ spl10_125 | ~ spl10_223), inference(avatar_contradiction_clause, [], [f12328])).
fof(f12328, plain, ($false | (spl10_105 | ~ spl10_125 | ~ spl10_223)), inference(subsumption_resolution, [], [f12320, f1082])).
fof(f1082, plain, (~ (e4 = op(e0, e4)) | spl10_105), inference(avatar_component_clause, [], [f1081])).
fof(f12320, plain, ((e4 = op(e0, e4)) | (~ spl10_125 | ~ spl10_223)), inference(backward_demodulation, [], [f1701, f1167])).
fof(f1701, plain, ((e4 = op(e0, op(e0, e0))) | ~ spl10_223), inference(avatar_component_clause, [], [f1700])).
fof(f12327, plain, (~ spl10_104 | ~ spl10_125 | spl10_222), inference(avatar_split_clause, [], [f12319, f1695, f1165, f1077])).
fof(f12319, plain, (~ (e3 = op(e0, e4)) | (~ spl10_125 | spl10_222)), inference(backward_demodulation, [], [f1697, f1167])).
fof(f12314, plain, (~ spl10_14 | ~ spl10_8 | spl10_168), inference(avatar_split_clause, [], [f12313, f1421, f674, f699])).
fof(f12313, plain, (~ (e3 = op(e4, e2)) | (~ spl10_8 | spl10_168)), inference(forward_demodulation, [], [f1423, f676])).
fof(f12270, plain, (~ spl10_34 | ~ spl10_44), inference(avatar_split_clause, [], [f12269, f825, f783])).
fof(f12269, plain, (~ (e3 = op(e3, e3)) | ~ spl10_44), inference(forward_demodulation, [], [f439, f827])).
fof(f827, plain, ((e3 = op(e3, e1)) | ~ spl10_44), inference(avatar_component_clause, [], [f825])).
fof(f12268, plain, (~ spl10_34 | ~ spl10_84), inference(avatar_split_clause, [], [f12267, f993, f783])).
fof(f12267, plain, (~ (e3 = op(e3, e3)) | ~ spl10_84), inference(forward_demodulation, [], [f389, f995])).
fof(f995, plain, ((e3 = op(e1, e3)) | ~ spl10_84), inference(avatar_component_clause, [], [f993])).
fof(f12264, plain, (~ spl10_32 | ~ spl10_84 | spl10_138), inference(avatar_split_clause, [], [f12130, f1273, f993, f775])).
fof(f12130, plain, (~ (e1 = op(e3, e3)) | (~ spl10_84 | spl10_138)), inference(backward_demodulation, [], [f1274, f995])).
fof(f12263, plain, (~ spl10_32 | ~ spl10_44 | spl10_178), inference(avatar_split_clause, [], [f12194, f1475, f825, f775])).
fof(f12194, plain, (~ (e1 = op(e3, e3)) | (~ spl10_44 | spl10_178)), inference(backward_demodulation, [], [f1477, f827])).
fof(f12260, plain, (~ spl10_5 | ~ spl10_20), inference(avatar_split_clause, [], [f12259, f724, f661])).
fof(f12259, plain, (~ (e4 = op(e4, e4)) | ~ spl10_20), inference(forward_demodulation, [], [f450, f726])).
fof(f726, plain, ((e4 = op(e4, e1)) | ~ spl10_20), inference(avatar_component_clause, [], [f724])).
fof(f12258, plain, (~ spl10_2 | ~ spl10_52), inference(avatar_split_clause, [], [f12257, f859, f649])).
fof(f12257, plain, (~ (e1 = op(e4, e4)) | ~ spl10_52), inference(forward_demodulation, [], [f402, f861])).
fof(f12255, plain, (~ spl10_2 | ~ spl10_20 | spl10_179), inference(avatar_split_clause, [], [f12216, f1481, f724, f649])).
fof(f12216, plain, (~ (e1 = op(e4, e4)) | (~ spl10_20 | spl10_179)), inference(backward_demodulation, [], [f1483, f726])).
fof(f1483, plain, (~ (e1 = op(e4, op(e4, e1))) | spl10_179), inference(avatar_component_clause, [], [f1481])).
fof(f12226, plain, (~ spl10_50 | ~ spl10_61 | ~ spl10_74 | spl10_212), inference(avatar_contradiction_clause, [], [f12225])).
fof(f12225, plain, ($false | (~ spl10_50 | ~ spl10_61 | ~ spl10_74 | spl10_212)), inference(subsumption_resolution, [], [f12224, f852])).
fof(f12224, plain, (~ (e4 = op(e3, e0)) | (~ spl10_61 | ~ spl10_74 | spl10_212)), inference(forward_demodulation, [], [f12223, f953])).
fof(f12223, plain, (~ (e4 = op(op(e2, e0), e0)) | (~ spl10_61 | spl10_212)), inference(forward_demodulation, [], [f1647, f899])).
fof(f1647, plain, (~ (e4 = op(op(e2, op(e2, e2)), op(e2, e2))) | spl10_212), inference(avatar_component_clause, [], [f1645])).
fof(f1645, plain, (spl10_212 <=> (e4 = op(op(e2, op(e2, e2)), op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl10_212])])).
fof(f12217, plain, (~ spl10_15 | ~ spl10_20), inference(avatar_split_clause, [], [f12215, f724, f703])).
fof(f12215, plain, (~ (e4 = op(e4, e2)) | ~ spl10_20), inference(backward_demodulation, [], [f448, f726])).
fof(f12208, plain, (~ spl10_37 | ~ spl10_68 | ~ spl10_170), inference(avatar_contradiction_clause, [], [f12207])).
fof(f12207, plain, ($false | (~ spl10_37 | ~ spl10_68 | ~ spl10_170)), inference(subsumption_resolution, [], [f12206, f461])).
fof(f12206, plain, ((e2 = e3) | (~ spl10_37 | ~ spl10_68 | ~ spl10_170)), inference(forward_demodulation, [], [f12203, f928])).
fof(f928, plain, ((e2 = op(e2, e1)) | ~ spl10_68), inference(avatar_component_clause, [], [f926])).
fof(f926, plain, (spl10_68 <=> (e2 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl10_68])])).
fof(f12203, plain, ((e3 = op(e2, e1)) | (~ spl10_37 | ~ spl10_170)), inference(backward_demodulation, [], [f1435, f798])).
fof(f12205, plain, (~ spl10_27 | ~ spl10_37), inference(avatar_split_clause, [], [f12199, f796, f754])).
fof(f12199, plain, (~ (e1 = op(e3, e4)) | ~ spl10_37), inference(backward_demodulation, [], [f442, f798])).
fof(f12196, plain, (~ spl10_39 | ~ spl10_44), inference(avatar_split_clause, [], [f12190, f825, f804])).
fof(f12190, plain, (~ (e3 = op(e3, e2)) | ~ spl10_44), inference(backward_demodulation, [], [f438, f827])).
fof(f12187, plain, (~ spl10_27 | ~ spl10_52), inference(avatar_split_clause, [], [f12181, f859, f754])).
fof(f12181, plain, (~ (e1 = op(e3, e4)) | ~ spl10_52), inference(backward_demodulation, [], [f401, f861])).
fof(f12174, plain, (~ spl10_61 | spl10_73 | ~ spl10_169), inference(avatar_contradiction_clause, [], [f12173])).
fof(f12173, plain, ($false | (~ spl10_61 | spl10_73 | ~ spl10_169)), inference(subsumption_resolution, [], [f12167, f948])).
fof(f948, plain, (~ (e2 = op(e2, e0)) | spl10_73), inference(avatar_component_clause, [], [f947])).
fof(f12167, plain, ((e2 = op(e2, e0)) | (~ spl10_61 | ~ spl10_169)), inference(backward_demodulation, [], [f1431, f899])).
fof(f12158, plain, (~ spl10_63 | ~ spl10_68), inference(avatar_split_clause, [], [f12150, f926, f905])).
fof(f12150, plain, (~ (e2 = op(e2, e2)) | ~ spl10_68), inference(backward_demodulation, [], [f428, f928])).
fof(f12157, plain, (~ spl10_43 | ~ spl10_68), inference(avatar_split_clause, [], [f12148, f926, f821])).
fof(f12148, plain, (~ (e2 = op(e3, e1)) | ~ spl10_68), inference(backward_demodulation, [], [f371, f928])).
fof(f12138, plain, (~ spl10_55 | ~ spl10_80), inference(avatar_split_clause, [], [f12134, f976, f871])).
fof(f12134, plain, (~ (e4 = op(e2, e4)) | ~ spl10_80), inference(backward_demodulation, [], [f398, f978])).
fof(f978, plain, ((e4 = op(e1, e4)) | ~ spl10_80), inference(avatar_component_clause, [], [f976])).
fof(f12132, plain, (~ spl10_59 | ~ spl10_84), inference(avatar_split_clause, [], [f12127, f993, f888])).
fof(f12127, plain, (~ (e3 = op(e2, e3)) | ~ spl10_84), inference(backward_demodulation, [], [f388, f995])).
fof(f388, plain, ~ (op(e1, e3) = op(e2, e3)), inference(cnf_transformation, [], [f4])).
fof(f12126, plain, (~ spl10_78 | ~ spl10_88), inference(avatar_split_clause, [], [f12120, f1010, f968])).
fof(f12120, plain, (~ (e2 = op(e1, e4)) | ~ spl10_88), inference(backward_demodulation, [], [f422, f1012])).
fof(f1012, plain, ((e2 = op(e1, e2)) | ~ spl10_88), inference(avatar_component_clause, [], [f1010])).
fof(f12125, plain, (~ spl10_63 | ~ spl10_88), inference(avatar_split_clause, [], [f12116, f1010, f905])).
fof(f12116, plain, (~ (e2 = op(e2, e2)) | ~ spl10_88), inference(backward_demodulation, [], [f378, f1012])).
fof(f12106, plain, (~ spl10_87 | ~ spl10_92), inference(avatar_split_clause, [], [f12090, f1027, f1006])).
fof(f12090, plain, (~ (e1 = op(e1, e2)) | ~ spl10_92), inference(backward_demodulation, [], [f418, f1029])).
fof(f1029, plain, ((e1 = op(e1, e1)) | ~ spl10_92), inference(avatar_component_clause, [], [f1027])).
fof(f12105, plain, (~ spl10_67 | ~ spl10_92), inference(avatar_split_clause, [], [f12087, f1027, f922])).
fof(f12087, plain, (~ (e1 = op(e2, e1)) | ~ spl10_92), inference(backward_demodulation, [], [f368, f1029])).
fof(f12086, plain, (~ spl10_91 | ~ spl10_96), inference(avatar_split_clause, [], [f12075, f1044, f1023])).
fof(f1044, plain, (spl10_96 <=> (e0 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl10_96])])).
fof(f12075, plain, (~ (e0 = op(e1, e1)) | ~ spl10_96), inference(backward_demodulation, [], [f414, f1046])).
fof(f1046, plain, ((e0 = op(e1, e0)) | ~ spl10_96), inference(avatar_component_clause, [], [f1044])).
fof(f12085, plain, (~ spl10_71 | ~ spl10_96), inference(avatar_split_clause, [], [f12074, f1044, f939])).
fof(f12074, plain, (~ (e0 = op(e2, e0)) | ~ spl10_96), inference(backward_demodulation, [], [f358, f1046])).
fof(f12072, plain, (spl10_103 | ~ spl10_115 | ~ spl10_148), inference(avatar_contradiction_clause, [], [f12071])).
fof(f12071, plain, ($false | (spl10_103 | ~ spl10_115 | ~ spl10_148)), inference(subsumption_resolution, [], [f12066, f1074])).
fof(f1074, plain, (~ (e2 = op(e0, e4)) | spl10_103), inference(avatar_component_clause, [], [f1073])).
fof(f12066, plain, ((e2 = op(e0, e4)) | (~ spl10_115 | ~ spl10_148)), inference(backward_demodulation, [], [f1321, f1125])).
fof(f1321, plain, ((e2 = op(e0, op(e0, e2))) | ~ spl10_148), inference(avatar_component_clause, [], [f1320])).
fof(f12059, plain, (~ spl10_111 | ~ spl10_116), inference(avatar_split_clause, [], [f12051, f1128, f1107])).
fof(f1128, plain, (spl10_116 <=> (e0 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl10_116])])).
fof(f12051, plain, (~ (e0 = op(e0, e2)) | ~ spl10_116), inference(backward_demodulation, [], [f408, f1130])).
fof(f1130, plain, ((e0 = op(e0, e1)) | ~ spl10_116), inference(avatar_component_clause, [], [f1128])).
fof(f12046, plain, (spl10_20 | ~ spl10_127), inference(avatar_split_clause, [], [f12033, f1174, f724])).
fof(f1174, plain, (spl10_127 <=> (e1 = unit)), introduced(avatar_definition, [new_symbols(naming, [spl10_127])])).
fof(f12033, plain, ((e4 = op(e4, e1)) | ~ spl10_127), inference(backward_demodulation, [], [f302, f1176])).
fof(f1176, plain, ((e1 = unit) | ~ spl10_127), inference(avatar_component_clause, [], [f1174])).
fof(f12045, plain, (spl10_80 | ~ spl10_127), inference(avatar_split_clause, [], [f12032, f1174, f976])).
fof(f12032, plain, ((e4 = op(e1, e4)) | ~ spl10_127), inference(backward_demodulation, [], [f301, f1176])).
fof(f12044, plain, (spl10_44 | ~ spl10_127), inference(avatar_split_clause, [], [f12031, f1174, f825])).
fof(f12031, plain, ((e3 = op(e3, e1)) | ~ spl10_127), inference(backward_demodulation, [], [f300, f1176])).
fof(f12043, plain, (spl10_84 | ~ spl10_127), inference(avatar_split_clause, [], [f12030, f1174, f993])).
fof(f12030, plain, ((e3 = op(e1, e3)) | ~ spl10_127), inference(backward_demodulation, [], [f299, f1176])).
fof(f12042, plain, (spl10_68 | ~ spl10_127), inference(avatar_split_clause, [], [f12029, f1174, f926])).
fof(f12029, plain, ((e2 = op(e2, e1)) | ~ spl10_127), inference(backward_demodulation, [], [f298, f1176])).
fof(f12041, plain, (spl10_88 | ~ spl10_127), inference(avatar_split_clause, [], [f12028, f1174, f1010])).
fof(f12028, plain, ((e2 = op(e1, e2)) | ~ spl10_127), inference(backward_demodulation, [], [f297, f1176])).
fof(f12040, plain, (spl10_92 | ~ spl10_127), inference(avatar_split_clause, [], [f12027, f1174, f1027])).
fof(f12027, plain, ((e1 = op(e1, e1)) | ~ spl10_127), inference(backward_demodulation, [], [f296, f1176])).
fof(f12038, plain, (spl10_116 | ~ spl10_127), inference(avatar_split_clause, [], [f12025, f1174, f1128])).
fof(f12025, plain, ((e0 = op(e0, e1)) | ~ spl10_127), inference(backward_demodulation, [], [f294, f1176])).
fof(f12037, plain, (spl10_96 | ~ spl10_127), inference(avatar_split_clause, [], [f12024, f1174, f1044])).
fof(f12024, plain, ((e0 = op(e1, e0)) | ~ spl10_127), inference(backward_demodulation, [], [f293, f1176])).
fof(f12014, plain, (spl10_115 | ~ spl10_123 | ~ spl10_223), inference(avatar_split_clause, [], [f12013, f1700, f1157, f1123])).
fof(f12013, plain, ((e4 = op(e0, e2)) | (~ spl10_123 | ~ spl10_223)), inference(forward_demodulation, [], [f1701, f1159])).
fof(f12007, plain, (~ spl10_22 | ~ spl10_31 | ~ spl10_50 | spl10_191), inference(avatar_contradiction_clause, [], [f12006])).
fof(f12006, plain, ($false | (~ spl10_22 | ~ spl10_31 | ~ spl10_50 | spl10_191)), inference(subsumption_resolution, [], [f12005, f735])).
fof(f12005, plain, (~ (e1 = op(e4, e0)) | (~ spl10_31 | ~ spl10_50 | spl10_191)), inference(forward_demodulation, [], [f12004, f852])).
fof(f12004, plain, (~ (e1 = op(op(e3, e0), e0)) | (~ spl10_31 | spl10_191)), inference(forward_demodulation, [], [f1545, f773])).
fof(f11990, plain, (~ spl10_113 | ~ spl10_123), inference(avatar_split_clause, [], [f11989, f1157, f1115])).
fof(f11989, plain, (~ (e2 = op(e0, e2)) | ~ spl10_123), inference(forward_demodulation, [], [f405, f1159])).
fof(f11988, plain, (~ spl10_114 | ~ spl10_123 | spl10_222), inference(avatar_split_clause, [], [f11875, f1695, f1157, f1119])).
fof(f11875, plain, (~ (e3 = op(e0, e2)) | (~ spl10_123 | spl10_222)), inference(backward_demodulation, [], [f1697, f1159])).
fof(f11986, plain, (~ spl10_112 | ~ spl10_123 | spl10_236), inference(avatar_split_clause, [], [f11880, f1781, f1157, f1111])).
fof(f11880, plain, (~ (e1 = op(e0, e2)) | (~ spl10_123 | spl10_236)), inference(backward_demodulation, [], [f1783, f1159])).
fof(f11981, plain, (~ spl10_73 | ~ spl10_123), inference(avatar_split_clause, [], [f11980, f1157, f947])).
fof(f11980, plain, (~ (e2 = op(e2, e0)) | ~ spl10_123), inference(forward_demodulation, [], [f355, f1159])).
fof(f11957, plain, (~ spl10_31 | spl10_121 | ~ spl10_190), inference(avatar_contradiction_clause, [], [f11956])).
fof(f11956, plain, ($false | (~ spl10_31 | spl10_121 | ~ spl10_190)), inference(subsumption_resolution, [], [f11950, f1150])).
fof(f1150, plain, (~ (e0 = op(e0, e0)) | spl10_121), inference(avatar_component_clause, [], [f1149])).
fof(f11950, plain, ((e0 = op(e0, e0)) | (~ spl10_31 | ~ spl10_190)), inference(backward_demodulation, [], [f1540, f773])).
fof(f11922, plain, (~ spl10_50 | ~ spl10_91 | ~ spl10_200 | spl10_211), inference(avatar_split_clause, [], [f11917, f1639, f1583, f1023, f850])).
fof(f11917, plain, (~ (e4 = op(e3, e0)) | (~ spl10_91 | ~ spl10_200 | spl10_211)), inference(backward_demodulation, [], [f11856, f1025])).
fof(f11856, plain, (~ (e4 = op(e3, op(e1, e1))) | (~ spl10_200 | spl10_211)), inference(backward_demodulation, [], [f1641, f1584])).
fof(f11921, plain, (~ spl10_48 | ~ spl10_91 | spl10_199 | ~ spl10_200), inference(avatar_split_clause, [], [f11916, f1583, f1579, f1023, f842])).
fof(f11916, plain, (~ (e2 = op(e3, e0)) | (~ spl10_91 | spl10_199 | ~ spl10_200)), inference(backward_demodulation, [], [f11854, f1025])).
fof(f11854, plain, (~ (e2 = op(e3, op(e1, e1))) | (spl10_199 | ~ spl10_200)), inference(backward_demodulation, [], [f1581, f1584])).
fof(f11919, plain, (~ spl10_41 | ~ spl10_91), inference(avatar_split_clause, [], [f11905, f1023, f813])).
fof(f11905, plain, (~ (e0 = op(e3, e1)) | ~ spl10_91), inference(backward_demodulation, [], [f369, f1025])).
fof(f11903, plain, (~ spl10_79 | ~ spl10_99), inference(avatar_split_clause, [], [f11900, f1056, f972])).
fof(f11900, plain, (~ (e3 = op(e1, e4)) | ~ spl10_99), inference(backward_demodulation, [], [f417, f1058])).
fof(f11882, plain, (~ spl10_103 | ~ spl10_123), inference(avatar_split_clause, [], [f11870, f1157, f1073])).
fof(f11870, plain, (~ (e2 = op(e0, e4)) | ~ spl10_123), inference(backward_demodulation, [], [f407, f1159])).
fof(f11867, plain, (~ spl10_41 | ~ spl10_107 | spl10_141), inference(avatar_split_clause, [], [f11866, f1287, f1090, f813])).
fof(f11866, plain, (~ (e0 = op(e3, e1)) | (~ spl10_107 | spl10_141)), inference(forward_demodulation, [], [f1288, f1092])).
fof(f11865, plain, (~ spl10_93 | ~ spl10_87 | spl10_145), inference(avatar_split_clause, [], [f11864, f1306, f1006, f1031])).
fof(f11864, plain, (~ (e2 = op(e1, e1)) | (~ spl10_87 | spl10_145)), inference(forward_demodulation, [], [f1308, f1008])).
fof(f11859, plain, (spl10_123 | ~ spl10_71 | ~ spl10_180), inference(avatar_split_clause, [], [f11858, f1489, f939, f1157])).
fof(f11858, plain, ((op(e0, e0) = e2) | (~ spl10_71 | ~ spl10_180)), inference(forward_demodulation, [], [f1491, f941])).
fof(f11849, plain, (~ spl10_124 | ~ spl10_1 | spl10_240), inference(avatar_split_clause, [], [f11848, f1813, f645, f1161])).
fof(f11848, plain, (~ (op(e0, e0) = e3) | (~ spl10_1 | spl10_240)), inference(forward_demodulation, [], [f1815, f647])).
fof(f11839, plain, (~ spl10_99 | ~ spl10_1 | ~ spl10_205 | spl10_210), inference(avatar_split_clause, [], [f11813, f1633, f1607, f645, f1056])).
fof(f11813, plain, (~ (e3 = op(e1, e0)) | (~ spl10_1 | ~ spl10_205 | spl10_210)), inference(backward_demodulation, [], [f11559, f647])).
fof(f11559, plain, (~ (e3 = op(e1, op(e4, e4))) | (~ spl10_205 | spl10_210)), inference(backward_demodulation, [], [f1635, f1608])).
fof(f11830, plain, (~ spl10_49 | ~ spl10_39), inference(avatar_split_clause, [], [f11829, f804, f846])).
fof(f11829, plain, (~ (e3 = op(e3, e0)) | ~ spl10_39), inference(forward_demodulation, [], [f435, f806])).
fof(f11791, plain, (~ spl10_19 | ~ spl10_81 | ~ spl10_176), inference(avatar_contradiction_clause, [], [f11790])).
fof(f11790, plain, ($false | (~ spl10_19 | ~ spl10_81 | ~ spl10_176)), inference(subsumption_resolution, [], [f11789, f457])).
fof(f11789, plain, ((e0 = e4) | (~ spl10_19 | ~ spl10_81 | ~ spl10_176)), inference(forward_demodulation, [], [f11786, f983])).
fof(f11786, plain, ((e4 = op(e1, e3)) | (~ spl10_19 | ~ spl10_176)), inference(backward_demodulation, [], [f1467, f722])).
fof(f11781, plain, (~ spl10_14 | ~ spl10_39), inference(avatar_split_clause, [], [f11777, f804, f699])).
fof(f11777, plain, (~ (e3 = op(e4, e2)) | ~ spl10_39), inference(backward_demodulation, [], [f383, f806])).
fof(f383, plain, ~ (op(e3, e2) = op(e4, e2)), inference(cnf_transformation, [], [f4])).
fof(f11685, plain, (~ spl10_88 | ~ spl10_93), inference(avatar_split_clause, [], [f11670, f1031, f1010])).
fof(f11670, plain, (~ (e2 = op(e1, e2)) | ~ spl10_93), inference(backward_demodulation, [], [f418, f1033])).
fof(f11666, plain, (~ spl10_80 | ~ spl10_100), inference(avatar_split_clause, [], [f11663, f1060, f976])).
fof(f11663, plain, (~ (e4 = op(e1, e4)) | ~ spl10_100), inference(backward_demodulation, [], [f417, f1062])).
fof(f11649, plain, (spl10_41 | ~ spl10_107 | ~ spl10_141), inference(avatar_split_clause, [], [f11646, f1287, f1090, f813])).
fof(f11646, plain, ((e0 = op(e3, e1)) | (~ spl10_107 | ~ spl10_141)), inference(backward_demodulation, [], [f1289, f1092])).
fof(f11614, plain, (spl10_107 | ~ spl10_124 | ~ spl10_236), inference(avatar_split_clause, [], [f11606, f1781, f1161, f1090])).
fof(f11606, plain, ((e1 = op(e0, e3)) | (~ spl10_124 | ~ spl10_236)), inference(backward_demodulation, [], [f1782, f1163])).
fof(f11613, plain, (~ spl10_104 | ~ spl10_124), inference(avatar_split_clause, [], [f11602, f1161, f1077])).
fof(f11602, plain, (~ (e3 = op(e0, e4)) | ~ spl10_124), inference(backward_demodulation, [], [f407, f1163])).
fof(f11544, plain, (~ spl10_97 | ~ spl10_22), inference(avatar_split_clause, [], [f11543, f733, f1048])).
fof(f11543, plain, (~ (e1 = op(e1, e0)) | ~ spl10_22), inference(forward_demodulation, [], [f360, f735])).
fof(f11524, plain, (~ spl10_36 | ~ spl10_48 | spl10_184), inference(avatar_split_clause, [], [f11493, f1507, f842, f792])).
fof(f11493, plain, (~ (e0 = op(e3, e2)) | (~ spl10_48 | spl10_184)), inference(backward_demodulation, [], [f1509, f844])).
fof(f11522, plain, (~ spl10_7 | ~ spl10_22), inference(avatar_split_clause, [], [f11521, f733, f670])).
fof(f11521, plain, (~ (e1 = op(e4, e3)) | ~ spl10_22), inference(forward_demodulation, [], [f446, f735])).
fof(f11511, plain, (~ spl10_27 | ~ spl10_35 | spl10_203), inference(avatar_contradiction_clause, [], [f11510])).
fof(f11510, plain, ($false | (~ spl10_27 | ~ spl10_35 | spl10_203)), inference(subsumption_resolution, [], [f11502, f756])).
fof(f11502, plain, (~ (e1 = op(e3, e4)) | (~ spl10_35 | spl10_203)), inference(backward_demodulation, [], [f1599, f789])).
fof(f1599, plain, (~ (e1 = op(e3, op(e3, e3))) | spl10_203), inference(avatar_component_clause, [], [f1597])).
fof(f11461, plain, (~ spl10_62 | ~ spl10_112), inference(avatar_split_clause, [], [f11454, f1111, f901])).
fof(f11454, plain, (~ (e1 = op(e2, e2)) | ~ spl10_112), inference(backward_demodulation, [], [f375, f1113])).
fof(f11426, plain, (spl10_60 | ~ spl10_14 | ~ spl10_171), inference(avatar_split_clause, [], [f11425, f1437, f699, f892])).
fof(f11425, plain, ((e4 = op(e2, e3)) | (~ spl10_14 | ~ spl10_171)), inference(forward_demodulation, [], [f1439, f701])).
fof(f11424, plain, (~ spl10_62 | ~ spl10_68 | spl10_177), inference(avatar_split_clause, [], [f11423, f1470, f926, f901])).
fof(f11423, plain, (~ (e1 = op(e2, e2)) | (~ spl10_68 | spl10_177)), inference(forward_demodulation, [], [f1472, f928])).
fof(f11422, plain, (spl10_108 | ~ spl10_74 | ~ spl10_180), inference(avatar_split_clause, [], [f11421, f1489, f951, f1094])).
fof(f11421, plain, ((e2 = op(e0, e3)) | (~ spl10_74 | ~ spl10_180)), inference(forward_demodulation, [], [f1491, f953])).
fof(f11415, plain, (~ spl10_62 | ~ spl10_3 | spl10_218), inference(avatar_split_clause, [], [f11414, f1677, f653, f901])).
fof(f1677, plain, (spl10_218 <=> (e1 = op(op(e4, e4), op(e4, e4)))), introduced(avatar_definition, [new_symbols(naming, [spl10_218])])).
fof(f11414, plain, (~ (e1 = op(e2, e2)) | (~ spl10_3 | spl10_218)), inference(forward_demodulation, [], [f1679, f655])).
fof(f1679, plain, (~ (e1 = op(op(e4, e4), op(e4, e4))) | spl10_218), inference(avatar_component_clause, [], [f1677])).
fof(f11382, plain, (~ spl10_110 | ~ spl10_104 | spl10_136), inference(avatar_split_clause, [], [f11277, f1264, f1077, f1102])).
fof(f11277, plain, (~ (e4 = op(e0, e3)) | (~ spl10_104 | spl10_136)), inference(backward_demodulation, [], [f1266, f1079])).
fof(f11374, plain, (~ spl10_48 | ~ spl10_36 | spl10_172), inference(avatar_split_clause, [], [f11340, f1443, f792, f842])).
fof(f11340, plain, (~ (e2 = op(e3, e0)) | (~ spl10_36 | spl10_172)), inference(backward_demodulation, [], [f1445, f794])).
fof(f11359, plain, (~ spl10_12 | ~ spl10_22), inference(avatar_split_clause, [], [f11353, f733, f691])).
fof(f11353, plain, (~ (e1 = op(e4, e2)) | ~ spl10_22), inference(backward_demodulation, [], [f445, f735])).
fof(f11346, plain, (~ spl10_33 | ~ spl10_36 | ~ spl10_115 | spl10_213), inference(avatar_contradiction_clause, [], [f11345])).
fof(f11345, plain, ($false | (~ spl10_33 | ~ spl10_36 | ~ spl10_115 | spl10_213)), inference(subsumption_resolution, [], [f11343, f1125])).
fof(f11343, plain, (~ (e4 = op(e0, e2)) | (~ spl10_33 | ~ spl10_36 | spl10_213)), inference(backward_demodulation, [], [f11234, f794])).
fof(f11234, plain, (~ (e4 = op(op(e3, e2), e2)) | (~ spl10_33 | spl10_213)), inference(forward_demodulation, [], [f1653, f781])).
fof(f11325, plain, (~ spl10_26 | ~ spl10_51), inference(avatar_split_clause, [], [f11318, f855, f750])).
fof(f11318, plain, (~ (e0 = op(e3, e4)) | ~ spl10_51), inference(backward_demodulation, [], [f401, f857])).
fof(f857, plain, ((e0 = op(e2, e4)) | ~ spl10_51), inference(avatar_component_clause, [], [f855])).
fof(f11308, plain, (~ spl10_62 | ~ spl10_68 | spl10_169), inference(avatar_contradiction_clause, [], [f11307])).
fof(f11307, plain, ($false | (~ spl10_62 | ~ spl10_68 | spl10_169)), inference(subsumption_resolution, [], [f11297, f928])).
fof(f11297, plain, (~ (e2 = op(e2, e1)) | (~ spl10_62 | spl10_169)), inference(backward_demodulation, [], [f1430, f903])).
fof(f11306, plain, (~ spl10_57 | ~ spl10_62), inference(avatar_split_clause, [], [f11292, f901, f880])).
fof(f11292, plain, (~ (e1 = op(e2, e3)) | ~ spl10_62), inference(backward_demodulation, [], [f431, f903])).
fof(f11305, plain, (~ spl10_12 | ~ spl10_62), inference(avatar_split_clause, [], [f11291, f901, f691])).
fof(f11291, plain, (~ (e1 = op(e4, e2)) | ~ spl10_62), inference(backward_demodulation, [], [f382, f903])).
fof(f11288, plain, (~ spl10_54 | ~ spl10_74), inference(avatar_split_clause, [], [f11284, f951, f867])).
fof(f11284, plain, (~ (e3 = op(e2, e4)) | ~ spl10_74), inference(backward_demodulation, [], [f427, f953])).
fof(f11278, plain, (~ spl10_54 | ~ spl10_104), inference(avatar_split_clause, [], [f11273, f1077, f867])).
fof(f11273, plain, (~ (e3 = op(e2, e4)) | ~ spl10_104), inference(backward_demodulation, [], [f395, f1079])).
fof(f395, plain, ~ (op(e0, e4) = op(e2, e4)), inference(cnf_transformation, [], [f4])).
fof(f11271, plain, (~ spl10_102 | ~ spl10_107), inference(avatar_split_clause, [], [f11264, f1090, f1069])).
fof(f11264, plain, (~ (e1 = op(e0, e4)) | ~ spl10_107), inference(backward_demodulation, [], [f413, f1092])).
fof(f11262, plain, (~ spl10_12 | ~ spl10_115 | ~ spl10_123 | spl10_235), inference(avatar_split_clause, [], [f11255, f1771, f1157, f1123, f691])).
fof(f11255, plain, (~ (e1 = op(e4, e2)) | (~ spl10_115 | ~ spl10_123 | spl10_235)), inference(backward_demodulation, [], [f11217, f1125])).
fof(f11217, plain, (~ (e1 = op(op(e0, e2), e2)) | (~ spl10_123 | spl10_235)), inference(forward_demodulation, [], [f1773, f1159])).
fof(f11260, plain, (spl10_51 | ~ spl10_115 | ~ spl10_147), inference(avatar_split_clause, [], [f11252, f1315, f1123, f855])).
fof(f11252, plain, ((e0 = op(e2, e4)) | (~ spl10_115 | ~ spl10_147)), inference(backward_demodulation, [], [f1317, f1125])).
fof(f11258, plain, (~ spl10_40 | ~ spl10_115), inference(avatar_split_clause, [], [f11248, f1123, f808])).
fof(f11248, plain, (~ (e4 = op(e3, e2)) | ~ spl10_115), inference(backward_demodulation, [], [f376, f1125])).
fof(f11243, plain, (spl10_50 | ~ spl10_6 | ~ spl10_166), inference(avatar_split_clause, [], [f11242, f1409, f666, f850])).
fof(f1409, plain, (spl10_166 <=> (e4 = op(e3, op(e4, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl10_166])])).
fof(f11242, plain, ((e4 = op(e3, e0)) | (~ spl10_6 | ~ spl10_166)), inference(forward_demodulation, [], [f1411, f668])).
fof(f1411, plain, ((e4 = op(e3, op(e4, e3))) | ~ spl10_166), inference(avatar_component_clause, [], [f1409])).
fof(f11239, plain, (~ spl10_61 | ~ spl10_33 | spl10_190), inference(avatar_split_clause, [], [f11238, f1539, f779, f897])).
fof(f11238, plain, (~ (e0 = op(e2, e2)) | (~ spl10_33 | spl10_190)), inference(forward_demodulation, [], [f1541, f781])).
fof(f11230, plain, (spl10_62 | ~ spl10_123 | ~ spl10_220), inference(avatar_split_clause, [], [f11229, f1687, f1157, f901])).
fof(f11229, plain, ((e1 = op(e2, e2)) | (~ spl10_123 | ~ spl10_220)), inference(forward_demodulation, [], [f1688, f1159])).
fof(f1688, plain, ((e1 = op(op(e0, e0), op(e0, e0))) | ~ spl10_220), inference(avatar_component_clause, [], [f1687])).
fof(f11224, plain, (spl10_36 | ~ spl10_33 | ~ spl10_224), inference(avatar_split_clause, [], [f11223, f1705, f779, f792])).
fof(f11223, plain, ((e0 = op(e3, e2)) | (~ spl10_33 | ~ spl10_224)), inference(forward_demodulation, [], [f1706, f781])).
fof(f11214, plain, (~ spl10_62 | ~ spl10_33 | spl10_216), inference(avatar_split_clause, [], [f11213, f1667, f779, f901])).
fof(f11213, plain, (~ (e1 = op(e2, e2)) | (~ spl10_33 | spl10_216)), inference(forward_demodulation, [], [f1669, f781])).
fof(f11207, plain, (~ spl10_105 | ~ spl10_80), inference(avatar_split_clause, [], [f11206, f976, f1081])).
fof(f11206, plain, (~ (e4 = op(e0, e4)) | ~ spl10_80), inference(forward_demodulation, [], [f394, f978])).
fof(f11187, plain, (~ spl10_24 | ~ spl10_6 | spl10_168), inference(avatar_split_clause, [], [f11161, f1421, f666, f741])).
fof(f11161, plain, (~ (e3 = op(e4, e0)) | (~ spl10_6 | spl10_168)), inference(backward_demodulation, [], [f1423, f668])).
fof(f11178, plain, (~ spl10_3 | ~ spl10_61 | spl10_194), inference(avatar_contradiction_clause, [], [f11177])).
fof(f11177, plain, ($false | (~ spl10_3 | ~ spl10_61 | spl10_194)), inference(subsumption_resolution, [], [f11168, f899])).
fof(f11168, plain, (~ (e0 = op(e2, e2)) | (~ spl10_3 | spl10_194)), inference(backward_demodulation, [], [f1559, f655])).
fof(f1559, plain, (~ (e0 = op(op(e4, e4), op(e4, e4))) | spl10_194), inference(avatar_component_clause, [], [f1557])).
fof(f1557, plain, (spl10_194 <=> (e0 = op(op(e4, e4), op(e4, e4)))), introduced(avatar_definition, [new_symbols(naming, [spl10_194])])).
fof(f11153, plain, (~ spl10_7 | ~ spl10_12), inference(avatar_split_clause, [], [f11146, f691, f670])).
fof(f11146, plain, (~ (e1 = op(e4, e3)) | ~ spl10_12), inference(backward_demodulation, [], [f451, f693])).
fof(f11145, plain, (~ spl10_6 | ~ spl10_24 | spl10_185), inference(avatar_split_clause, [], [f11143, f1513, f741, f666])).
fof(f11143, plain, (~ (e0 = op(e4, e3)) | (~ spl10_24 | spl10_185)), inference(backward_demodulation, [], [f1515, f743])).
fof(f11125, plain, (~ spl10_12 | ~ spl10_33 | ~ spl10_40 | spl10_191), inference(avatar_split_clause, [], [f11124, f1543, f808, f779, f691])).
fof(f11124, plain, (~ (e1 = op(e4, e2)) | (~ spl10_33 | ~ spl10_40 | spl10_191)), inference(forward_demodulation, [], [f11117, f810])).
fof(f11117, plain, (~ (e1 = op(op(e3, e2), e2)) | (~ spl10_33 | spl10_191)), inference(backward_demodulation, [], [f1545, f781])).
fof(f11111, plain, (~ spl10_28 | ~ spl10_40 | spl10_172), inference(avatar_split_clause, [], [f11109, f1443, f808, f758])).
fof(f11109, plain, (~ (e2 = op(e3, e4)) | (~ spl10_40 | spl10_172)), inference(backward_demodulation, [], [f1445, f810])).
fof(f11103, plain, (~ spl10_7 | ~ spl10_57), inference(avatar_split_clause, [], [f11097, f880, f670])).
fof(f11097, plain, (~ (e1 = op(e4, e3)) | ~ spl10_57), inference(backward_demodulation, [], [f392, f882])).
fof(f11074, plain, (spl10_26 | ~ spl10_110 | ~ spl10_141), inference(avatar_split_clause, [], [f11071, f1287, f1102, f750])).
fof(f11071, plain, ((e0 = op(e3, e4)) | (~ spl10_110 | ~ spl10_141)), inference(backward_demodulation, [], [f1289, f1104])).
fof(f11066, plain, (~ spl10_56 | ~ spl10_114 | spl10_147), inference(avatar_split_clause, [], [f11063, f1315, f1119, f876])).
fof(f11063, plain, (~ (e0 = op(e2, e3)) | (~ spl10_114 | spl10_147)), inference(backward_demodulation, [], [f1316, f1121])).
fof(f1316, plain, (~ (e0 = op(e2, op(e0, e2))) | spl10_147), inference(avatar_component_clause, [], [f1315])).
fof(f11021, plain, (~ spl10_222 | ~ spl10_236), inference(avatar_contradiction_clause, [], [f11020])).
fof(f11020, plain, ($false | (~ spl10_222 | ~ spl10_236)), inference(subsumption_resolution, [], [f11019, f459])).
fof(f459, plain, ~ (e1 = e3), inference(cnf_transformation, [], [f5])).
fof(f11019, plain, ((e1 = e3) | (~ spl10_222 | ~ spl10_236)), inference(forward_demodulation, [], [f1696, f1782])).
fof(f11001, plain, (~ spl10_29 | ~ spl10_44), inference(avatar_split_clause, [], [f11000, f825, f762])).
fof(f11000, plain, (~ (e3 = op(e3, e4)) | ~ spl10_44), inference(forward_demodulation, [], [f440, f827])).
fof(f10999, plain, (~ spl10_30 | ~ spl10_80), inference(avatar_split_clause, [], [f10998, f976, f766])).
fof(f10998, plain, (~ (e4 = op(e3, e4)) | ~ spl10_80), inference(forward_demodulation, [], [f399, f978])).
fof(f10997, plain, (~ spl10_29 | ~ spl10_54), inference(avatar_split_clause, [], [f10996, f867, f762])).
fof(f10996, plain, (~ (e3 = op(e3, e4)) | ~ spl10_54), inference(forward_demodulation, [], [f401, f869])).
fof(f10993, plain, (~ spl10_25 | ~ spl10_75), inference(avatar_split_clause, [], [f10992, f955, f745])).
fof(f10992, plain, (~ (e4 = op(e4, e0)) | ~ spl10_75), inference(forward_demodulation, [], [f362, f957])).
fof(f10991, plain, (~ spl10_10 | ~ spl10_20), inference(avatar_split_clause, [], [f10990, f724, f682])).
fof(f10990, plain, (~ (e4 = op(e4, e3)) | ~ spl10_20), inference(forward_demodulation, [], [f449, f726])).
fof(f10989, plain, (~ spl10_9 | ~ spl10_84), inference(avatar_split_clause, [], [f10988, f993, f678])).
fof(f10988, plain, (~ (e3 = op(e4, e3)) | ~ spl10_84), inference(forward_demodulation, [], [f390, f995])).
fof(f10967, plain, (~ spl10_14 | ~ spl10_56 | ~ spl10_171), inference(avatar_contradiction_clause, [], [f10966])).
fof(f10966, plain, ($false | (~ spl10_14 | ~ spl10_56 | ~ spl10_171)), inference(subsumption_resolution, [], [f10965, f457])).
fof(f10965, plain, ((e0 = e4) | (~ spl10_14 | ~ spl10_56 | ~ spl10_171)), inference(forward_demodulation, [], [f10962, f878])).
fof(f10962, plain, ((e4 = op(e2, e3)) | (~ spl10_14 | ~ spl10_171)), inference(backward_demodulation, [], [f1439, f701])).
fof(f10952, plain, (~ spl10_28 | ~ spl10_35 | spl10_192), inference(avatar_contradiction_clause, [], [f10951])).
fof(f10951, plain, ($false | (~ spl10_28 | ~ spl10_35 | spl10_192)), inference(subsumption_resolution, [], [f10943, f760])).
fof(f10943, plain, (~ (e2 = op(e3, e4)) | (~ spl10_35 | spl10_192)), inference(backward_demodulation, [], [f1549, f789])).
fof(f1549, plain, (~ (e2 = op(e3, op(e3, e3))) | spl10_192), inference(avatar_component_clause, [], [f1547])).
fof(f10921, plain, (~ spl10_4 | ~ spl10_54), inference(avatar_split_clause, [], [f10918, f867, f657])).
fof(f10918, plain, (~ (e3 = op(e4, e4)) | ~ spl10_54), inference(backward_demodulation, [], [f402, f869])).
fof(f10879, plain, (~ spl10_55 | ~ spl10_75), inference(avatar_split_clause, [], [f10876, f955, f871])).
fof(f10876, plain, (~ (e4 = op(e2, e4)) | ~ spl10_75), inference(backward_demodulation, [], [f427, f957])).
fof(f10818, plain, (~ spl10_76 | ~ spl10_96), inference(avatar_split_clause, [], [f10810, f1044, f960])).
fof(f10810, plain, (~ (e0 = op(e1, e4)) | ~ spl10_96), inference(backward_demodulation, [], [f417, f1046])).
fof(f10797, plain, (~ spl10_16 | ~ spl10_116), inference(avatar_split_clause, [], [f10789, f1128, f708])).
fof(f10789, plain, (~ (e0 = op(e4, e1)) | ~ spl10_116), inference(backward_demodulation, [], [f367, f1130])).
fof(f10785, plain, (~ spl10_85 | ~ spl10_124 | spl10_229 | ~ spl10_236), inference(avatar_split_clause, [], [f10777, f1781, f1735, f1161, f997])).
fof(f1735, plain, (spl10_229 <=> (e4 = op(op(e0, op(e0, e0)), op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl10_229])])).
fof(f10777, plain, (~ (e4 = op(e1, e3)) | (~ spl10_124 | spl10_229 | ~ spl10_236)), inference(backward_demodulation, [], [f10601, f1163])).
fof(f10601, plain, (~ (e4 = op(e1, op(e0, e0))) | (spl10_229 | ~ spl10_236)), inference(backward_demodulation, [], [f1737, f1782])).
fof(f1737, plain, (~ (e4 = op(op(e0, op(e0, e0)), op(e0, e0))) | spl10_229), inference(avatar_component_clause, [], [f1735])).
fof(f10743, plain, (~ spl10_16 | ~ spl10_102 | spl10_135), inference(avatar_split_clause, [], [f10742, f1259, f1069, f708])).
fof(f10742, plain, (~ (e0 = op(e4, e1)) | (~ spl10_102 | spl10_135)), inference(forward_demodulation, [], [f1260, f1071])).
fof(f10733, plain, (spl10_115 | ~ spl10_23 | ~ spl10_182), inference(avatar_split_clause, [], [f10732, f1497, f737, f1123])).
fof(f10732, plain, ((e4 = op(e0, e2)) | (~ spl10_23 | ~ spl10_182)), inference(forward_demodulation, [], [f1499, f739])).
fof(f10731, plain, (~ spl10_190 | ~ spl10_244), inference(avatar_contradiction_clause, [], [f10730])).
fof(f10730, plain, ($false | (~ spl10_190 | ~ spl10_244)), inference(subsumption_resolution, [], [f10729, f457])).
fof(f10729, plain, ((e0 = e4) | (~ spl10_190 | ~ spl10_244)), inference(backward_demodulation, [], [f1854, f1540])).
fof(f1854, plain, ((e4 = op(op(e3, e3), op(e3, e3))) | ~ spl10_244), inference(avatar_component_clause, [], [f1853])).
fof(f10723, plain, (~ spl10_227 | ~ spl10_236), inference(avatar_contradiction_clause, [], [f10722])).
fof(f10722, plain, ($false | (~ spl10_227 | ~ spl10_236)), inference(subsumption_resolution, [], [f10721, f458])).
fof(f458, plain, ~ (e1 = e2), inference(cnf_transformation, [], [f5])).
fof(f10721, plain, ((e1 = e2) | (~ spl10_227 | ~ spl10_236)), inference(forward_demodulation, [], [f1722, f1782])).
fof(f10708, plain, (~ spl10_40 | ~ spl10_28 | spl10_163), inference(avatar_split_clause, [], [f10392, f1391, f758, f808])).
fof(f10392, plain, (~ (e4 = op(e3, e2)) | (~ spl10_28 | spl10_163)), inference(forward_demodulation, [], [f1393, f760])).
fof(f10698, plain, (~ spl10_16 | ~ spl10_99 | ~ spl10_176), inference(avatar_contradiction_clause, [], [f10697])).
fof(f10697, plain, ($false | (~ spl10_16 | ~ spl10_99 | ~ spl10_176)), inference(subsumption_resolution, [], [f10696, f463])).
fof(f463, plain, ~ (e3 = e4), inference(cnf_transformation, [], [f5])).
fof(f10696, plain, ((e3 = e4) | (~ spl10_16 | ~ spl10_99 | ~ spl10_176)), inference(forward_demodulation, [], [f10693, f1058])).
fof(f10693, plain, ((e4 = op(e1, e0)) | (~ spl10_16 | ~ spl10_176)), inference(backward_demodulation, [], [f1467, f710])).
fof(f10695, plain, (~ spl10_1 | ~ spl10_16), inference(avatar_split_clause, [], [f10688, f708, f645])).
fof(f10688, plain, (~ (e0 = op(e4, e4)) | ~ spl10_16), inference(backward_demodulation, [], [f450, f710])).
fof(f10635, plain, (~ spl10_100 | ~ spl10_125), inference(avatar_split_clause, [], [f10627, f1165, f1060])).
fof(f10627, plain, (~ (e4 = op(e1, e0)) | ~ spl10_125), inference(backward_demodulation, [], [f354, f1167])).
fof(f10624, plain, (~ spl10_120 | ~ spl10_102 | spl10_136), inference(avatar_split_clause, [], [f10623, f1264, f1069, f1144])).
fof(f10623, plain, (~ (e4 = op(e0, e1)) | (~ spl10_102 | spl10_136)), inference(forward_demodulation, [], [f1266, f1071])).
fof(f10622, plain, (spl10_45 | ~ spl10_7 | ~ spl10_166), inference(avatar_split_clause, [], [f10621, f1409, f670, f829])).
fof(f10621, plain, ((e4 = op(e3, e1)) | (~ spl10_7 | ~ spl10_166)), inference(forward_demodulation, [], [f1411, f672])).
fof(f672, plain, ((e1 = op(e4, e3)) | ~ spl10_7), inference(avatar_component_clause, [], [f670])).
fof(f10620, plain, (~ spl10_19 | ~ spl10_7 | spl10_168), inference(avatar_split_clause, [], [f10619, f1421, f670, f720])).
fof(f10619, plain, (~ (e3 = op(e4, e1)) | (~ spl10_7 | spl10_168)), inference(forward_demodulation, [], [f1423, f672])).
fof(f10617, plain, (spl10_119 | ~ spl10_47 | ~ spl10_181), inference(avatar_split_clause, [], [f10616, f1493, f838, f1140])).
fof(f10616, plain, ((e3 = op(e0, e1)) | (~ spl10_47 | ~ spl10_181)), inference(forward_demodulation, [], [f1495, f840])).
fof(f10598, plain, (~ spl10_241 | ~ spl10_245), inference(avatar_contradiction_clause, [], [f10597])).
fof(f10597, plain, ($false | (~ spl10_241 | ~ spl10_245)), inference(subsumption_resolution, [], [f10596, f463])).
fof(f10596, plain, ((e3 = e4) | (~ spl10_241 | ~ spl10_245)), inference(backward_demodulation, [], [f1860, f1820])).
fof(f10572, plain, (~ spl10_3 | ~ spl10_23), inference(avatar_split_clause, [], [f10571, f737, f653])).
fof(f10571, plain, (~ (e2 = op(e4, e4)) | ~ spl10_23), inference(forward_demodulation, [], [f447, f739])).
fof(f10547, plain, (~ spl10_41 | ~ spl10_47 | spl10_178), inference(avatar_contradiction_clause, [], [f10546])).
fof(f10546, plain, ($false | (~ spl10_41 | ~ spl10_47 | spl10_178)), inference(subsumption_resolution, [], [f10542, f840])).
fof(f10542, plain, (~ (e1 = op(e3, e0)) | (~ spl10_41 | spl10_178)), inference(backward_demodulation, [], [f1477, f815])).
fof(f10543, plain, (~ spl10_31 | ~ spl10_41), inference(avatar_split_clause, [], [f10536, f813, f771])).
fof(f10536, plain, (~ (e0 = op(e3, e3)) | ~ spl10_41), inference(backward_demodulation, [], [f439, f815])).
fof(f10535, plain, (~ spl10_31 | ~ spl10_81), inference(avatar_split_clause, [], [f10528, f981, f771])).
fof(f10528, plain, (~ (e0 = op(e3, e3)) | ~ spl10_81), inference(backward_demodulation, [], [f389, f983])).
fof(f10520, plain, (~ spl10_83 | ~ spl10_93), inference(avatar_split_clause, [], [f10509, f1031, f989])).
fof(f10509, plain, (~ (e2 = op(e1, e3)) | ~ spl10_93), inference(backward_demodulation, [], [f419, f1033])).
fof(f10476, plain, (~ spl10_110 | ~ spl10_124 | spl10_223), inference(avatar_split_clause, [], [f10475, f1700, f1161, f1102])).
fof(f10475, plain, (~ (e4 = op(e0, e3)) | (~ spl10_124 | spl10_223)), inference(forward_demodulation, [], [f1702, f1163])).
fof(f10464, plain, (~ spl10_108 | ~ spl10_124 | spl10_227), inference(avatar_split_clause, [], [f10402, f1721, f1161, f1094])).
fof(f10402, plain, (~ (e2 = op(e0, e3)) | (~ spl10_124 | spl10_227)), inference(backward_demodulation, [], [f1723, f1163])).
fof(f10453, plain, (~ spl10_85 | ~ spl10_79 | spl10_133), inference(avatar_split_clause, [], [f10174, f1250, f972, f997])).
fof(f1250, plain, (spl10_133 <=> (e4 = op(e1, op(e1, e4)))), introduced(avatar_definition, [new_symbols(naming, [spl10_133])])).
fof(f10174, plain, (~ (e4 = op(e1, e3)) | (~ spl10_79 | spl10_133)), inference(backward_demodulation, [], [f1252, f974])).
fof(f1252, plain, (~ (e4 = op(e1, op(e1, e4))) | spl10_133), inference(avatar_component_clause, [], [f1250])).
fof(f10452, plain, (~ spl10_85 | ~ spl10_19 | spl10_176), inference(avatar_split_clause, [], [f10439, f1465, f720, f997])).
fof(f10439, plain, (~ (e4 = op(e1, e3)) | (~ spl10_19 | spl10_176)), inference(backward_demodulation, [], [f1466, f722])).
fof(f1466, plain, (~ (e4 = op(e1, op(e4, e1))) | spl10_176), inference(avatar_component_clause, [], [f1465])).
fof(f10444, plain, (~ spl10_32 | ~ spl10_124 | spl10_220), inference(avatar_split_clause, [], [f10399, f1687, f1161, f775])).
fof(f10399, plain, (~ (e1 = op(e3, e3)) | (~ spl10_124 | spl10_220)), inference(backward_demodulation, [], [f1689, f1163])).
fof(f10443, plain, (~ spl10_7 | ~ spl10_19 | spl10_179), inference(avatar_contradiction_clause, [], [f10442])).
fof(f10442, plain, ($false | (~ spl10_7 | ~ spl10_19 | spl10_179)), inference(subsumption_resolution, [], [f10440, f672])).
fof(f10440, plain, (~ (e1 = op(e4, e3)) | (~ spl10_19 | spl10_179)), inference(backward_demodulation, [], [f1483, f722])).
fof(f10432, plain, (~ spl10_18 | ~ spl10_23), inference(avatar_split_clause, [], [f10424, f737, f716])).
fof(f10424, plain, (~ (e2 = op(e4, e1)) | ~ spl10_23), inference(backward_demodulation, [], [f444, f739])).
fof(f10413, plain, (~ spl10_7 | ~ spl10_124 | ~ spl10_223 | spl10_235), inference(avatar_contradiction_clause, [], [f10412])).
fof(f10412, plain, ($false | (~ spl10_7 | ~ spl10_124 | ~ spl10_223 | spl10_235)), inference(subsumption_resolution, [], [f10404, f672])).
fof(f10404, plain, (~ (e1 = op(e4, e3)) | (~ spl10_124 | ~ spl10_223 | spl10_235)), inference(backward_demodulation, [], [f10376, f1163])).
fof(f10376, plain, (~ (e1 = op(e4, op(e0, e0))) | (~ spl10_223 | spl10_235)), inference(backward_demodulation, [], [f1773, f1701])).
fof(f10411, plain, (~ spl10_31 | ~ spl10_124 | ~ spl10_234), inference(avatar_contradiction_clause, [], [f10410])).
fof(f10410, plain, ($false | (~ spl10_31 | ~ spl10_124 | ~ spl10_234)), inference(subsumption_resolution, [], [f10409, f455])).
fof(f455, plain, ~ (e0 = e2), inference(cnf_transformation, [], [f5])).
fof(f10409, plain, ((e0 = e2) | (~ spl10_31 | ~ spl10_124 | ~ spl10_234)), inference(forward_demodulation, [], [f10403, f773])).
fof(f10403, plain, ((e2 = op(e3, e3)) | (~ spl10_124 | ~ spl10_234)), inference(backward_demodulation, [], [f1768, f1163])).
fof(f10406, plain, (~ spl10_24 | ~ spl10_124), inference(avatar_split_clause, [], [f10396, f1161, f741])).
fof(f10396, plain, (~ (e3 = op(e4, e0)) | ~ spl10_124), inference(backward_demodulation, [], [f357, f1163])).
fof(f10395, plain, (~ spl10_123 | ~ spl10_111 | spl10_148), inference(avatar_split_clause, [], [f10394, f1320, f1107, f1157])).
fof(f10394, plain, (~ (op(e0, e0) = e2) | (~ spl10_111 | spl10_148)), inference(forward_demodulation, [], [f1322, f1109])).
fof(f10374, plain, (~ spl10_123 | ~ spl10_91 | spl10_230), inference(avatar_split_clause, [], [f10373, f1745, f1023, f1157])).
fof(f10373, plain, (~ (op(e0, e0) = e2) | (~ spl10_91 | spl10_230)), inference(forward_demodulation, [], [f1747, f1025])).
fof(f10364, plain, (spl10_24 | ~ spl10_1 | ~ spl10_197), inference(avatar_split_clause, [], [f10363, f1570, f645, f741])).
fof(f1570, plain, (spl10_197 <=> (e3 = op(e4, op(e4, e4)))), introduced(avatar_definition, [new_symbols(naming, [spl10_197])])).
fof(f10363, plain, ((e3 = op(e4, e0)) | (~ spl10_1 | ~ spl10_197)), inference(forward_demodulation, [], [f1571, f647])).
fof(f1571, plain, ((e3 = op(e4, op(e4, e4))) | ~ spl10_197), inference(avatar_component_clause, [], [f1570])).
fof(f10362, plain, (~ spl10_123 | ~ spl10_1 | spl10_233), inference(avatar_split_clause, [], [f10361, f1761, f645, f1157])).
fof(f10361, plain, (~ (op(e0, e0) = e2) | (~ spl10_1 | spl10_233)), inference(forward_demodulation, [], [f1763, f647])).
fof(f10355, plain, (~ spl10_125 | ~ spl10_110), inference(avatar_split_clause, [], [f10354, f1102, f1165])).
fof(f10354, plain, (~ (op(e0, e0) = e4) | ~ spl10_110), inference(forward_demodulation, [], [f406, f1104])).
fof(f10351, plain, (~ spl10_123 | ~ spl10_31 | spl10_232), inference(avatar_split_clause, [], [f10273, f1755, f771, f1157])).
fof(f1755, plain, (spl10_232 <=> (e2 = op(op(e3, e3), op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl10_232])])).
fof(f10273, plain, (~ (op(e0, e0) = e2) | (~ spl10_31 | spl10_232)), inference(backward_demodulation, [], [f1757, f773])).
fof(f1757, plain, (~ (e2 = op(op(e3, e3), op(e3, e3))) | spl10_232), inference(avatar_component_clause, [], [f1755])).
fof(f10350, plain, (~ spl10_125 | ~ spl10_31 | spl10_244), inference(avatar_split_clause, [], [f10274, f1853, f771, f1165])).
fof(f10274, plain, (~ (op(e0, e0) = e4) | (~ spl10_31 | spl10_244)), inference(backward_demodulation, [], [f1855, f773])).
fof(f10313, plain, (~ spl10_3 | ~ spl10_15 | spl10_173), inference(avatar_split_clause, [], [f10311, f1449, f703, f653])).
fof(f10311, plain, (~ (e2 = op(e4, e4)) | (~ spl10_15 | spl10_173)), inference(backward_demodulation, [], [f1451, f705])).
fof(f10288, plain, (~ spl10_3 | ~ spl10_28), inference(avatar_split_clause, [], [f10283, f758, f653])).
fof(f10283, plain, (~ (e2 = op(e4, e4)) | ~ spl10_28), inference(backward_demodulation, [], [f403, f760])).
fof(f10282, plain, (~ spl10_31 | ~ spl10_100 | ~ spl10_203 | spl10_213), inference(avatar_contradiction_clause, [], [f10281])).
fof(f10281, plain, ($false | (~ spl10_31 | ~ spl10_100 | ~ spl10_203 | spl10_213)), inference(subsumption_resolution, [], [f10276, f1062])).
fof(f10276, plain, (~ (e4 = op(e1, e0)) | (~ spl10_31 | ~ spl10_203 | spl10_213)), inference(backward_demodulation, [], [f10035, f773])).
fof(f10035, plain, (~ (e4 = op(e1, op(e3, e3))) | (~ spl10_203 | spl10_213)), inference(backward_demodulation, [], [f1653, f1598])).
fof(f10245, plain, (~ spl10_22 | ~ spl10_47), inference(avatar_split_clause, [], [f10236, f838, f733])).
fof(f10236, plain, (~ (e1 = op(e4, e0)) | ~ spl10_47), inference(backward_demodulation, [], [f363, f840])).
fof(f10176, plain, (spl10_7 | ~ spl10_79 | ~ spl10_132), inference(avatar_split_clause, [], [f10173, f1245, f972, f670])).
fof(f10173, plain, ((e1 = op(e4, e3)) | (~ spl10_79 | ~ spl10_132)), inference(backward_demodulation, [], [f1247, f974])).
fof(f10167, plain, (~ spl10_58 | ~ spl10_83), inference(avatar_split_clause, [], [f10159, f989, f884])).
fof(f10159, plain, (~ (e2 = op(e2, e3)) | ~ spl10_83), inference(backward_demodulation, [], [f388, f991])).
fof(f10144, plain, (~ spl10_24 | ~ spl10_91 | ~ spl10_201 | spl10_206), inference(avatar_split_clause, [], [f10139, f1613, f1588, f1023, f741])).
fof(f1613, plain, (spl10_206 <=> (e3 = op(op(e1, op(e1, e1)), op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl10_206])])).
fof(f10139, plain, (~ (e3 = op(e4, e0)) | (~ spl10_91 | ~ spl10_201 | spl10_206)), inference(backward_demodulation, [], [f10038, f1025])).
fof(f10038, plain, (~ (e3 = op(e4, op(e1, e1))) | (~ spl10_201 | spl10_206)), inference(backward_demodulation, [], [f1615, f1589])).
fof(f1589, plain, ((e4 = op(e1, op(e1, e1))) | ~ spl10_201), inference(avatar_component_clause, [], [f1588])).
fof(f1615, plain, (~ (e3 = op(op(e1, op(e1, e1)), op(e1, e1))) | spl10_206), inference(avatar_component_clause, [], [f1613])).
fof(f10143, plain, (~ spl10_76 | ~ spl10_91), inference(avatar_split_clause, [], [f10128, f1023, f960])).
fof(f10128, plain, (~ (e0 = op(e1, e4)) | ~ spl10_91), inference(backward_demodulation, [], [f420, f1025])).
fof(f10142, plain, (~ spl10_66 | ~ spl10_91), inference(avatar_split_clause, [], [f10123, f1023, f918])).
fof(f10123, plain, (~ (e0 = op(e2, e1)) | ~ spl10_91), inference(backward_demodulation, [], [f368, f1025])).
fof(f10122, plain, (~ spl10_76 | ~ spl10_100 | spl10_156), inference(avatar_split_clause, [], [f10119, f1358, f1060, f960])).
fof(f10119, plain, (~ (e0 = op(e1, e4)) | (~ spl10_100 | spl10_156)), inference(backward_demodulation, [], [f1360, f1062])).
fof(f10109, plain, (~ spl10_10 | ~ spl10_110), inference(avatar_split_clause, [], [f10106, f1102, f682])).
fof(f10106, plain, (~ (e4 = op(e4, e3)) | ~ spl10_110), inference(backward_demodulation, [], [f387, f1104])).
fof(f10009, plain, (~ spl10_19 | ~ spl10_82 | ~ spl10_176), inference(avatar_contradiction_clause, [], [f10008])).
fof(f10008, plain, ($false | (~ spl10_19 | ~ spl10_82 | ~ spl10_176)), inference(subsumption_resolution, [], [f10007, f460])).
fof(f460, plain, ~ (e1 = e4), inference(cnf_transformation, [], [f5])).
fof(f10007, plain, ((e1 = e4) | (~ spl10_19 | ~ spl10_82 | ~ spl10_176)), inference(forward_demodulation, [], [f10006, f987])).
fof(f10006, plain, ((e4 = op(e1, e3)) | (~ spl10_19 | ~ spl10_176)), inference(forward_demodulation, [], [f1467, f722])).
fof(f10004, plain, (~ spl10_3 | ~ spl10_11 | ~ spl10_114 | spl10_210), inference(avatar_contradiction_clause, [], [f10003])).
fof(f10003, plain, ($false | (~ spl10_3 | ~ spl10_11 | ~ spl10_114 | spl10_210)), inference(subsumption_resolution, [], [f10002, f1121])).
fof(f10002, plain, (~ (e3 = op(e0, e2)) | (~ spl10_3 | ~ spl10_11 | spl10_210)), inference(forward_demodulation, [], [f10001, f689])).
fof(f10001, plain, (~ (e3 = op(op(e4, e2), e2)) | (~ spl10_3 | spl10_210)), inference(forward_demodulation, [], [f1635, f655])).
fof(f9977, plain, (~ spl10_11 | ~ spl10_90 | ~ spl10_93 | spl10_231), inference(avatar_contradiction_clause, [], [f9976])).
fof(f9976, plain, ($false | (~ spl10_11 | ~ spl10_90 | ~ spl10_93 | spl10_231)), inference(subsumption_resolution, [], [f9975, f689])).
fof(f9975, plain, (~ (e0 = op(e4, e2)) | (~ spl10_90 | ~ spl10_93 | spl10_231)), inference(forward_demodulation, [], [f9974, f1020])).
fof(f9974, plain, (~ (e0 = op(op(e1, e2), e2)) | (~ spl10_93 | spl10_231)), inference(forward_demodulation, [], [f1751, f1033])).
fof(f9964, plain, (~ spl10_34 | spl10_165), inference(avatar_contradiction_clause, [], [f9963])).
fof(f9963, plain, ($false | (~ spl10_34 | spl10_165)), inference(subsumption_resolution, [], [f9962, f785])).
fof(f785, plain, ((e3 = op(e3, e3)) | ~ spl10_34), inference(avatar_component_clause, [], [f783])).
fof(f9962, plain, (~ (e3 = op(e3, e3)) | (~ spl10_34 | spl10_165)), inference(forward_demodulation, [], [f1406, f785])).
fof(f9959, plain, (~ spl10_62 | ~ spl10_66 | ~ spl10_120 | spl10_212), inference(avatar_contradiction_clause, [], [f9958])).
fof(f9958, plain, ($false | (~ spl10_62 | ~ spl10_66 | ~ spl10_120 | spl10_212)), inference(subsumption_resolution, [], [f9957, f1146])).
fof(f9957, plain, (~ (e4 = op(e0, e1)) | (~ spl10_62 | ~ spl10_66 | spl10_212)), inference(forward_demodulation, [], [f9956, f920])).
fof(f9956, plain, (~ (e4 = op(op(e2, e1), e1)) | (~ spl10_62 | spl10_212)), inference(forward_demodulation, [], [f1647, f903])).
fof(f9931, plain, (~ spl10_6 | ~ spl10_11), inference(avatar_split_clause, [], [f9923, f687, f666])).
fof(f9923, plain, (~ (e0 = op(e4, e3)) | ~ spl10_11), inference(backward_demodulation, [], [f451, f689])).
fof(f9921, plain, (~ spl10_4 | ~ spl10_19), inference(avatar_split_clause, [], [f9917, f720, f657])).
fof(f9917, plain, (~ (e3 = op(e4, e4)) | ~ spl10_19), inference(backward_demodulation, [], [f450, f722])).
fof(f9882, plain, (~ spl10_42 | spl10_178), inference(avatar_contradiction_clause, [], [f9881])).
fof(f9881, plain, ($false | (~ spl10_42 | spl10_178)), inference(subsumption_resolution, [], [f9876, f819])).
fof(f9876, plain, (~ (e1 = op(e3, e1)) | (~ spl10_42 | spl10_178)), inference(backward_demodulation, [], [f1477, f819])).
fof(f9880, plain, (~ spl10_42 | ~ spl10_93 | ~ spl10_175), inference(avatar_contradiction_clause, [], [f9879])).
fof(f9879, plain, ($false | (~ spl10_42 | ~ spl10_93 | ~ spl10_175)), inference(subsumption_resolution, [], [f9878, f461])).
fof(f9878, plain, ((e2 = e3) | (~ spl10_42 | ~ spl10_93 | ~ spl10_175)), inference(forward_demodulation, [], [f9875, f1033])).
fof(f9875, plain, ((e3 = op(e1, e1)) | (~ spl10_42 | ~ spl10_175)), inference(backward_demodulation, [], [f1463, f819])).
fof(f9877, plain, (~ spl10_32 | ~ spl10_42), inference(avatar_split_clause, [], [f9870, f817, f775])).
fof(f9870, plain, (~ (e1 = op(e3, e3)) | ~ spl10_42), inference(backward_demodulation, [], [f439, f819])).
fof(f9867, plain, (~ spl10_41 | ~ spl10_46), inference(avatar_split_clause, [], [f9858, f834, f813])).
fof(f9858, plain, (~ (e0 = op(e3, e1)) | ~ spl10_46), inference(backward_demodulation, [], [f434, f836])).
fof(f836, plain, ((e0 = op(e3, e0)) | ~ spl10_46), inference(avatar_component_clause, [], [f834])).
fof(f9838, plain, (~ spl10_62 | ~ spl10_66 | spl10_228), inference(avatar_contradiction_clause, [], [f9837])).
fof(f9837, plain, ($false | (~ spl10_62 | ~ spl10_66 | spl10_228)), inference(subsumption_resolution, [], [f9831, f920])).
fof(f9831, plain, (~ (e0 = op(e2, e1)) | (~ spl10_62 | spl10_228)), inference(backward_demodulation, [], [f1729, f903])).
fof(f9835, plain, (~ spl10_62 | ~ spl10_66 | ~ spl10_169), inference(avatar_contradiction_clause, [], [f9834])).
fof(f9834, plain, ($false | (~ spl10_62 | ~ spl10_66 | ~ spl10_169)), inference(subsumption_resolution, [], [f9833, f455])).
fof(f9833, plain, ((e0 = e2) | (~ spl10_62 | ~ spl10_66 | ~ spl10_169)), inference(forward_demodulation, [], [f9827, f920])).
fof(f9827, plain, ((e2 = op(e2, e1)) | (~ spl10_62 | ~ spl10_169)), inference(backward_demodulation, [], [f1431, f903])).
fof(f9819, plain, (~ spl10_41 | ~ spl10_66), inference(avatar_split_clause, [], [f9808, f918, f813])).
fof(f9808, plain, (~ (e0 = op(e3, e1)) | ~ spl10_66), inference(backward_demodulation, [], [f371, f920])).
fof(f9795, plain, (~ spl10_32 | ~ spl10_82), inference(avatar_split_clause, [], [f9788, f985, f775])).
fof(f9788, plain, (~ (e1 = op(e3, e3)) | ~ spl10_82), inference(backward_demodulation, [], [f389, f987])).
fof(f9775, plain, (spl10_86 | ~ spl10_93 | ~ spl10_237), inference(avatar_contradiction_clause, [], [f9774])).
fof(f9774, plain, ($false | (spl10_86 | ~ spl10_93 | ~ spl10_237)), inference(subsumption_resolution, [], [f9766, f1003])).
fof(f1003, plain, (~ (e0 = op(e1, e2)) | spl10_86), inference(avatar_component_clause, [], [f1002])).
fof(f9766, plain, ((e0 = op(e1, e2)) | (~ spl10_93 | ~ spl10_237)), inference(backward_demodulation, [], [f1788, f1033])).
fof(f9771, plain, (~ spl10_18 | ~ spl10_93), inference(avatar_split_clause, [], [f9759, f1031, f716])).
fof(f9759, plain, (~ (e2 = op(e4, e1)) | ~ spl10_93), inference(backward_demodulation, [], [f370, f1033])).
fof(f9754, plain, (~ spl10_6 | ~ spl10_106), inference(avatar_split_clause, [], [f9747, f1086, f666])).
fof(f9747, plain, (~ (e0 = op(e4, e3)) | ~ spl10_106), inference(backward_demodulation, [], [f387, f1088])).
fof(f1088, plain, ((e0 = op(e0, e3)) | ~ spl10_106), inference(avatar_component_clause, [], [f1086])).
fof(f9668, plain, (~ spl10_209 | ~ spl10_169 | spl10_187), inference(avatar_split_clause, [], [f9667, f1525, f1429, f1627])).
fof(f9667, plain, (~ (e1 = op(e2, op(e2, e2))) | (~ spl10_169 | spl10_187)), inference(forward_demodulation, [], [f1527, f1431])).
fof(f9666, plain, (~ spl10_64 | ~ spl10_123 | spl10_241), inference(avatar_split_clause, [], [f9544, f1819, f1157, f909])).
fof(f9544, plain, (~ (e3 = op(e2, e2)) | (~ spl10_123 | spl10_241)), inference(backward_demodulation, [], [f1821, f1159])).
fof(f9665, plain, (~ spl10_65 | ~ spl10_123 | spl10_245), inference(avatar_split_clause, [], [f9545, f1859, f1157, f913])).
fof(f9545, plain, (~ (e4 = op(e2, e2)) | (~ spl10_123 | spl10_245)), inference(backward_demodulation, [], [f1861, f1159])).
fof(f9628, plain, (~ spl10_26 | ~ spl10_41), inference(avatar_split_clause, [], [f9617, f813, f750])).
fof(f9617, plain, (~ (e0 = op(e3, e4)) | ~ spl10_41), inference(backward_demodulation, [], [f440, f815])).
fof(f9627, plain, (~ spl10_16 | ~ spl10_41), inference(avatar_split_clause, [], [f9616, f813, f708])).
fof(f9616, plain, (~ (e0 = op(e4, e1)) | ~ spl10_41), inference(backward_demodulation, [], [f373, f815])).
fof(f9614, plain, (~ spl10_50 | ~ spl10_102 | ~ spl10_181), inference(avatar_contradiction_clause, [], [f9613])).
fof(f9613, plain, ($false | (~ spl10_50 | ~ spl10_102 | ~ spl10_181)), inference(subsumption_resolution, [], [f9612, f459])).
fof(f9612, plain, ((e1 = e3) | (~ spl10_50 | ~ spl10_102 | ~ spl10_181)), inference(forward_demodulation, [], [f9609, f1071])).
fof(f9609, plain, ((e3 = op(e0, e4)) | (~ spl10_50 | ~ spl10_181)), inference(backward_demodulation, [], [f1495, f852])).
fof(f9611, plain, (~ spl10_45 | ~ spl10_50), inference(avatar_split_clause, [], [f9607, f850, f829])).
fof(f9607, plain, (~ (e4 = op(e3, e1)) | ~ spl10_50), inference(backward_demodulation, [], [f434, f852])).
fof(f9605, plain, (~ spl10_26 | ~ spl10_76), inference(avatar_split_clause, [], [f9598, f960, f750])).
fof(f9598, plain, (~ (e0 = op(e3, e4)) | ~ spl10_76), inference(backward_demodulation, [], [f399, f962])).
fof(f9596, plain, (~ spl10_8 | ~ spl10_83), inference(avatar_split_clause, [], [f9591, f989, f674])).
fof(f9591, plain, (~ (e2 = op(e4, e3)) | ~ spl10_83), inference(backward_demodulation, [], [f390, f991])).
fof(f9553, plain, (~ spl10_111 | ~ spl10_123 | ~ spl10_223), inference(avatar_contradiction_clause, [], [f9552])).
fof(f9552, plain, ($false | (~ spl10_111 | ~ spl10_123 | ~ spl10_223)), inference(subsumption_resolution, [], [f9551, f457])).
fof(f9551, plain, ((e0 = e4) | (~ spl10_111 | ~ spl10_123 | ~ spl10_223)), inference(forward_demodulation, [], [f9541, f1109])).
fof(f9541, plain, ((e4 = op(e0, e2)) | (~ spl10_123 | ~ spl10_223)), inference(backward_demodulation, [], [f1701, f1159])).
fof(f9523, plain, (~ spl10_16 | ~ spl10_22 | spl10_185), inference(avatar_split_clause, [], [f9522, f1513, f733, f708])).
fof(f9522, plain, (~ (e0 = op(e4, e1)) | (~ spl10_22 | spl10_185)), inference(forward_demodulation, [], [f1515, f735])).
fof(f9519, plain, (~ spl10_8 | ~ spl10_4 | spl10_196), inference(avatar_split_clause, [], [f9518, f1565, f657, f674])).
fof(f9518, plain, (~ (e2 = op(e4, e3)) | (~ spl10_4 | spl10_196)), inference(forward_demodulation, [], [f1567, f659])).
fof(f659, plain, ((e3 = op(e4, e4)) | ~ spl10_4), inference(avatar_component_clause, [], [f657])).
fof(f9514, plain, (spl10_41 | ~ spl10_32 | ~ spl10_224), inference(avatar_split_clause, [], [f9513, f1705, f775, f813])).
fof(f9513, plain, ((e0 = op(e3, e1)) | (~ spl10_32 | ~ spl10_224)), inference(forward_demodulation, [], [f1706, f777])).
fof(f9462, plain, (~ spl10_16 | ~ spl10_22 | spl10_179), inference(avatar_contradiction_clause, [], [f9461])).
fof(f9461, plain, ($false | (~ spl10_16 | ~ spl10_22 | spl10_179)), inference(subsumption_resolution, [], [f9459, f735])).
fof(f9459, plain, (~ (e1 = op(e4, e0)) | (~ spl10_16 | spl10_179)), inference(backward_demodulation, [], [f1483, f710])).
fof(f9440, plain, (~ spl10_16 | ~ spl10_32 | ~ spl10_45 | spl10_217), inference(avatar_split_clause, [], [f9439, f1671, f829, f775, f708])).
fof(f9439, plain, (~ (e0 = op(e4, e1)) | (~ spl10_32 | ~ spl10_45 | spl10_217)), inference(forward_demodulation, [], [f9430, f831])).
fof(f9430, plain, (~ (e0 = op(op(e3, e1), e1)) | (~ spl10_32 | spl10_217)), inference(backward_demodulation, [], [f1673, f777])).
fof(f9436, plain, (~ spl10_32 | ~ spl10_45 | spl10_193), inference(avatar_contradiction_clause, [], [f9435])).
fof(f9435, plain, ($false | (~ spl10_32 | ~ spl10_45 | spl10_193)), inference(subsumption_resolution, [], [f9426, f831])).
fof(f9426, plain, (~ (e4 = op(e3, e1)) | (~ spl10_32 | spl10_193)), inference(backward_demodulation, [], [f1554, f777])).
fof(f9369, plain, (~ spl10_63 | spl10_169), inference(avatar_contradiction_clause, [], [f9368])).
fof(f9368, plain, ($false | (~ spl10_63 | spl10_169)), inference(subsumption_resolution, [], [f9360, f907])).
fof(f9360, plain, (~ (e2 = op(e2, e2)) | (~ spl10_63 | spl10_169)), inference(backward_demodulation, [], [f1430, f907])).
fof(f9305, plain, (~ spl10_94 | ~ spl10_110 | spl10_211 | ~ spl10_237), inference(avatar_contradiction_clause, [], [f9304])).
fof(f9304, plain, ($false | (~ spl10_94 | ~ spl10_110 | spl10_211 | ~ spl10_237)), inference(subsumption_resolution, [], [f9299, f1104])).
fof(f9299, plain, (~ (e4 = op(e0, e3)) | (~ spl10_94 | spl10_211 | ~ spl10_237)), inference(backward_demodulation, [], [f9190, f1037])).
fof(f9190, plain, (~ (e4 = op(e0, op(e1, e1))) | (spl10_211 | ~ spl10_237)), inference(backward_demodulation, [], [f1641, f1788])).
fof(f9302, plain, (spl10_81 | ~ spl10_94 | ~ spl10_237), inference(avatar_split_clause, [], [f9295, f1787, f1035, f981])).
fof(f9295, plain, ((e0 = op(e1, e3)) | (~ spl10_94 | ~ spl10_237)), inference(backward_demodulation, [], [f1788, f1037])).
fof(f9241, plain, (~ spl10_8 | ~ spl10_124 | spl10_221 | ~ spl10_223), inference(avatar_split_clause, [], [f9236, f1700, f1691, f1161, f674])).
fof(f9236, plain, (~ (e2 = op(e4, e3)) | (~ spl10_124 | spl10_221 | ~ spl10_223)), inference(backward_demodulation, [], [f9193, f1163])).
fof(f9193, plain, (~ (e2 = op(e4, op(e0, e0))) | (spl10_221 | ~ spl10_223)), inference(forward_demodulation, [], [f1693, f1701])).
fof(f9240, plain, (spl10_110 | ~ spl10_124 | ~ spl10_223), inference(avatar_split_clause, [], [f9233, f1700, f1161, f1102])).
fof(f9233, plain, ((e4 = op(e0, e3)) | (~ spl10_124 | ~ spl10_223)), inference(backward_demodulation, [], [f1701, f1163])).
fof(f9199, plain, (spl10_16 | ~ spl10_22 | ~ spl10_185), inference(avatar_split_clause, [], [f9198, f1513, f733, f708])).
fof(f9198, plain, ((e0 = op(e4, e1)) | (~ spl10_22 | ~ spl10_185)), inference(forward_demodulation, [], [f1514, f735])).
fof(f1514, plain, ((e0 = op(e4, op(e4, e0))) | ~ spl10_185), inference(avatar_component_clause, [], [f1513])).
fof(f9175, plain, (~ spl10_52 | ~ spl10_102), inference(avatar_split_clause, [], [f9174, f1069, f859])).
fof(f9174, plain, (~ (e1 = op(e2, e4)) | ~ spl10_102), inference(forward_demodulation, [], [f395, f1071])).
fof(f9173, plain, (~ spl10_53 | ~ spl10_78), inference(avatar_split_clause, [], [f9172, f968, f863])).
fof(f9172, plain, (~ (e2 = op(e2, e4)) | ~ spl10_78), inference(forward_demodulation, [], [f398, f970])).
fof(f9155, plain, (~ spl10_12 | ~ spl10_18 | spl10_179), inference(avatar_split_clause, [], [f9154, f1481, f716, f691])).
fof(f9154, plain, (~ (e1 = op(e4, e2)) | (~ spl10_18 | spl10_179)), inference(backward_demodulation, [], [f1483, f718])).
fof(f9139, plain, (~ spl10_30 | spl10_163), inference(avatar_contradiction_clause, [], [f9138])).
fof(f9138, plain, ($false | (~ spl10_30 | spl10_163)), inference(subsumption_resolution, [], [f9137, f768])).
fof(f768, plain, ((e4 = op(e3, e4)) | ~ spl10_30), inference(avatar_component_clause, [], [f766])).
fof(f9137, plain, (~ (e4 = op(e3, e4)) | (~ spl10_30 | spl10_163)), inference(backward_demodulation, [], [f1393, f768])).
fof(f9103, plain, (~ spl10_36 | ~ spl10_46), inference(avatar_split_clause, [], [f9095, f834, f792])).
fof(f9095, plain, (~ (e0 = op(e3, e2)) | ~ spl10_46), inference(backward_demodulation, [], [f435, f836])).
fof(f9081, plain, (~ spl10_62 | ~ spl10_69 | spl10_188), inference(avatar_contradiction_clause, [], [f9080])).
fof(f9080, plain, ($false | (~ spl10_62 | ~ spl10_69 | spl10_188)), inference(subsumption_resolution, [], [f9072, f932])).
fof(f9072, plain, (~ (e3 = op(e2, e1)) | (~ spl10_62 | spl10_188)), inference(backward_demodulation, [], [f1531, f903])).
fof(f9025, plain, (spl10_78 | ~ spl10_90 | ~ spl10_145), inference(avatar_split_clause, [], [f9021, f1306, f1018, f968])).
fof(f9021, plain, ((e2 = op(e1, e4)) | (~ spl10_90 | ~ spl10_145)), inference(backward_demodulation, [], [f1307, f1020])).
fof(f1307, plain, ((e2 = op(e1, op(e1, e2))) | ~ spl10_145), inference(avatar_component_clause, [], [f1306])).
fof(f8988, plain, (~ spl10_84 | ~ spl10_99), inference(avatar_split_clause, [], [f8982, f1056, f993])).
fof(f8982, plain, (~ (e3 = op(e1, e3)) | ~ spl10_99), inference(backward_demodulation, [], [f416, f1058])).
fof(f8907, plain, (~ spl10_6 | ~ spl10_4 | spl10_225), inference(avatar_split_clause, [], [f8906, f1711, f657, f666])).
fof(f8906, plain, (~ (e0 = op(e4, e3)) | (~ spl10_4 | spl10_225)), inference(forward_demodulation, [], [f1713, f659])).
fof(f1713, plain, (~ (e0 = op(e4, op(e4, e4))) | spl10_225), inference(avatar_component_clause, [], [f1711])).
fof(f8900, plain, (~ spl10_109 | ~ spl10_114), inference(avatar_split_clause, [], [f8899, f1119, f1098])).
fof(f8899, plain, (~ (e3 = op(e0, e3)) | ~ spl10_114), inference(forward_demodulation, [], [f411, f1121])).
fof(f8865, plain, (~ spl10_65 | spl10_101 | ~ spl10_215 | ~ spl10_228), inference(avatar_contradiction_clause, [], [f8864])).
fof(f8864, plain, ($false | (~ spl10_65 | spl10_101 | ~ spl10_215 | ~ spl10_228)), inference(subsumption_resolution, [], [f8862, f1066])).
fof(f1066, plain, (~ (e0 = op(e0, e4)) | spl10_101), inference(avatar_component_clause, [], [f1065])).
fof(f8862, plain, ((e0 = op(e0, e4)) | (~ spl10_65 | ~ spl10_215 | ~ spl10_228)), inference(backward_demodulation, [], [f8807, f915])).
fof(f8807, plain, ((e0 = op(e0, op(e2, e2))) | (~ spl10_215 | ~ spl10_228)), inference(backward_demodulation, [], [f1662, f1728])).
fof(f1662, plain, ((e0 = op(op(e2, op(e2, e2)), op(e2, e2))) | ~ spl10_215), inference(avatar_component_clause, [], [f1661])).
fof(f8857, plain, (~ spl10_64 | ~ spl10_74), inference(avatar_split_clause, [], [f8852, f951, f909])).
fof(f8852, plain, (~ (e3 = op(e2, e2)) | ~ spl10_74), inference(backward_demodulation, [], [f425, f953])).
fof(f8845, plain, (~ spl10_102 | ~ spl10_125 | spl10_236), inference(avatar_contradiction_clause, [], [f8844])).
fof(f8844, plain, ($false | (~ spl10_102 | ~ spl10_125 | spl10_236)), inference(subsumption_resolution, [], [f8838, f1071])).
fof(f8838, plain, (~ (e1 = op(e0, e4)) | (~ spl10_125 | spl10_236)), inference(backward_demodulation, [], [f1783, f1167])).
fof(f8826, plain, (spl10_74 | ~ spl10_36 | ~ spl10_170), inference(avatar_split_clause, [], [f8825, f1433, f792, f951])).
fof(f8825, plain, ((e3 = op(e2, e0)) | (~ spl10_36 | ~ spl10_170)), inference(forward_demodulation, [], [f1435, f794])).
fof(f8821, plain, (~ spl10_115 | ~ spl10_23 | spl10_182), inference(avatar_split_clause, [], [f8820, f1497, f737, f1123])).
fof(f8820, plain, (~ (e4 = op(e0, e2)) | (~ spl10_23 | spl10_182)), inference(forward_demodulation, [], [f1498, f739])).
fof(f8809, plain, (~ spl10_32 | ~ spl10_4 | spl10_218), inference(avatar_split_clause, [], [f8808, f1677, f657, f775])).
fof(f8808, plain, (~ (e1 = op(e3, e3)) | (~ spl10_4 | spl10_218)), inference(forward_demodulation, [], [f1679, f659])).
fof(f8778, plain, (~ spl10_4 | ~ spl10_6 | ~ spl10_108 | spl10_204), inference(avatar_contradiction_clause, [], [f8777])).
fof(f8777, plain, ($false | (~ spl10_4 | ~ spl10_6 | ~ spl10_108 | spl10_204)), inference(subsumption_resolution, [], [f8773, f1096])).
fof(f8773, plain, (~ (e2 = op(e0, e3)) | (~ spl10_4 | ~ spl10_6 | spl10_204)), inference(backward_demodulation, [], [f7810, f668])).
fof(f7810, plain, (~ (e2 = op(op(e4, e3), e3)) | (~ spl10_4 | spl10_204)), inference(forward_demodulation, [], [f1605, f659])).
fof(f8759, plain, (~ spl10_11 | ~ spl10_23 | spl10_185), inference(avatar_split_clause, [], [f8757, f1513, f737, f687])).
fof(f8757, plain, (~ (e0 = op(e4, e2)) | (~ spl10_23 | spl10_185)), inference(backward_demodulation, [], [f1515, f739])).
fof(f8758, plain, (~ spl10_8 | ~ spl10_23), inference(avatar_split_clause, [], [f8753, f737, f674])).
fof(f8753, plain, (~ (e2 = op(e4, e3)) | ~ spl10_23), inference(backward_demodulation, [], [f446, f739])).
fof(f8746, plain, (~ spl10_32 | spl10_43 | ~ spl10_192), inference(avatar_contradiction_clause, [], [f8745])).
fof(f8745, plain, ($false | (~ spl10_32 | spl10_43 | ~ spl10_192)), inference(subsumption_resolution, [], [f8737, f822])).
fof(f822, plain, (~ (e2 = op(e3, e1)) | spl10_43), inference(avatar_component_clause, [], [f821])).
fof(f8737, plain, ((e2 = op(e3, e1)) | (~ spl10_32 | ~ spl10_192)), inference(backward_demodulation, [], [f1548, f777])).
fof(f8728, plain, (~ spl10_11 | ~ spl10_36), inference(avatar_split_clause, [], [f8720, f792, f687])).
fof(f8720, plain, (~ (e0 = op(e4, e2)) | ~ spl10_36), inference(backward_demodulation, [], [f383, f794])).
fof(f8711, plain, (~ spl10_60 | ~ spl10_64 | spl10_189), inference(avatar_contradiction_clause, [], [f8710])).
fof(f8710, plain, ($false | (~ spl10_60 | ~ spl10_64 | spl10_189)), inference(subsumption_resolution, [], [f8706, f894])).
fof(f8706, plain, (~ (e4 = op(e2, e3)) | (~ spl10_64 | spl10_189)), inference(backward_demodulation, [], [f1536, f911])).
fof(f1536, plain, (~ (e4 = op(e2, op(e2, e2))) | spl10_189), inference(avatar_component_clause, [], [f1534])).
fof(f1534, plain, (spl10_189 <=> (e4 = op(e2, op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl10_189])])).
fof(f8698, plain, (~ spl10_62 | ~ spl10_72), inference(avatar_split_clause, [], [f8692, f943, f901])).
fof(f8692, plain, (~ (e1 = op(e2, e2)) | ~ spl10_72), inference(backward_demodulation, [], [f425, f945])).
fof(f8697, plain, (~ spl10_22 | ~ spl10_72), inference(avatar_split_clause, [], [f8691, f943, f733])).
fof(f8691, plain, (~ (e1 = op(e4, e0)) | ~ spl10_72), inference(backward_demodulation, [], [f362, f945])).
fof(f8638, plain, (spl10_23 | ~ spl10_51 | ~ spl10_159), inference(avatar_split_clause, [], [f8637, f1373, f855, f737])).
fof(f8637, plain, ((e2 = op(e4, e0)) | (~ spl10_51 | ~ spl10_159)), inference(forward_demodulation, [], [f1375, f857])).
fof(f8634, plain, (~ spl10_165 | ~ spl10_192), inference(avatar_contradiction_clause, [], [f8633])).
fof(f8633, plain, ($false | (~ spl10_165 | ~ spl10_192)), inference(subsumption_resolution, [], [f8632, f461])).
fof(f8632, plain, ((e2 = e3) | (~ spl10_165 | ~ spl10_192)), inference(forward_demodulation, [], [f1407, f1548])).
fof(f1407, plain, ((e3 = op(e3, op(e3, e3))) | ~ spl10_165), inference(avatar_component_clause, [], [f1405])).
fof(f8626, plain, (~ spl10_31 | ~ spl10_4 | spl10_194), inference(avatar_split_clause, [], [f8625, f1557, f657, f771])).
fof(f8625, plain, (~ (e0 = op(e3, e3)) | (~ spl10_4 | spl10_194)), inference(forward_demodulation, [], [f1559, f659])).
fof(f8594, plain, (~ spl10_75 | ~ spl10_51 | spl10_162), inference(avatar_split_clause, [], [f8523, f1386, f855, f955])).
fof(f8523, plain, (~ (e4 = op(e2, e0)) | (~ spl10_51 | spl10_162)), inference(backward_demodulation, [], [f1388, f857])).
fof(f8581, plain, (~ spl10_11 | ~ spl10_74 | ~ spl10_171), inference(avatar_contradiction_clause, [], [f8580])).
fof(f8580, plain, ($false | (~ spl10_11 | ~ spl10_74 | ~ spl10_171)), inference(subsumption_resolution, [], [f8579, f463])).
fof(f8579, plain, ((e3 = e4) | (~ spl10_11 | ~ spl10_74 | ~ spl10_171)), inference(forward_demodulation, [], [f8577, f953])).
fof(f8577, plain, ((e4 = op(e2, e0)) | (~ spl10_11 | ~ spl10_171)), inference(backward_demodulation, [], [f1439, f689])).
fof(f8552, plain, (~ spl10_26 | ~ spl10_31), inference(avatar_split_clause, [], [f8539, f771, f750])).
fof(f8539, plain, (~ (e0 = op(e3, e4)) | ~ spl10_31), inference(backward_demodulation, [], [f443, f773])).
fof(f8511, plain, (~ spl10_62 | spl10_91 | ~ spl10_186), inference(avatar_contradiction_clause, [], [f8510])).
fof(f8510, plain, ($false | (~ spl10_62 | spl10_91 | ~ spl10_186)), inference(subsumption_resolution, [], [f8503, f1024])).
fof(f1024, plain, (~ (e0 = op(e1, e1)) | spl10_91), inference(avatar_component_clause, [], [f1023])).
fof(f8503, plain, ((e0 = op(e1, e1)) | (~ spl10_62 | ~ spl10_186)), inference(backward_demodulation, [], [f1522, f903])).
fof(f8485, plain, (~ spl10_49 | ~ spl10_74), inference(avatar_split_clause, [], [f8478, f951, f846])).
fof(f8478, plain, (~ (e3 = op(e3, e0)) | ~ spl10_74), inference(backward_demodulation, [], [f361, f953])).
fof(f8425, plain, (~ spl10_81 | ~ spl10_96), inference(avatar_split_clause, [], [f8417, f1044, f981])).
fof(f8417, plain, (~ (e0 = op(e1, e3)) | ~ spl10_96), inference(backward_demodulation, [], [f416, f1046])).
fof(f8389, plain, (spl10_108 | ~ spl10_114 | ~ spl10_148), inference(avatar_contradiction_clause, [], [f8388])).
fof(f8388, plain, ($false | (spl10_108 | ~ spl10_114 | ~ spl10_148)), inference(subsumption_resolution, [], [f8385, f1095])).
fof(f1095, plain, (~ (e2 = op(e0, e3)) | spl10_108), inference(avatar_component_clause, [], [f1094])).
fof(f8385, plain, ((e2 = op(e0, e3)) | (~ spl10_114 | ~ spl10_148)), inference(backward_demodulation, [], [f1321, f1121])).
fof(f8290, plain, (~ spl10_33 | ~ spl10_4 | spl10_233), inference(avatar_split_clause, [], [f8153, f1761, f657, f779])).
fof(f8153, plain, (~ (e2 = op(e3, e3)) | (~ spl10_4 | spl10_233)), inference(forward_demodulation, [], [f1763, f659])).
fof(f8284, plain, (~ spl10_40 | ~ spl10_52 | ~ spl10_170), inference(avatar_contradiction_clause, [], [f8283])).
fof(f8283, plain, ($false | (~ spl10_40 | ~ spl10_52 | ~ spl10_170)), inference(subsumption_resolution, [], [f8282, f459])).
fof(f8282, plain, ((e1 = e3) | (~ spl10_40 | ~ spl10_52 | ~ spl10_170)), inference(forward_demodulation, [], [f8281, f861])).
fof(f8281, plain, ((e3 = op(e2, e4)) | (~ spl10_40 | ~ spl10_170)), inference(forward_demodulation, [], [f1435, f810])).
fof(f8277, plain, (~ spl10_32 | ~ spl10_43 | spl10_192), inference(avatar_contradiction_clause, [], [f8276])).
fof(f8276, plain, ($false | (~ spl10_32 | ~ spl10_43 | spl10_192)), inference(subsumption_resolution, [], [f8275, f823])).
fof(f8275, plain, (~ (e2 = op(e3, e1)) | (~ spl10_32 | spl10_192)), inference(forward_demodulation, [], [f1549, f777])).
fof(f8246, plain, (~ spl10_61 | ~ spl10_73 | spl10_169), inference(avatar_contradiction_clause, [], [f8245])).
fof(f8245, plain, ($false | (~ spl10_61 | ~ spl10_73 | spl10_169)), inference(subsumption_resolution, [], [f8244, f949])).
fof(f8244, plain, (~ (e2 = op(e2, e0)) | (~ spl10_61 | spl10_169)), inference(forward_demodulation, [], [f1430, f899])).
fof(f8222, plain, (~ spl10_61 | ~ spl10_121 | spl10_186), inference(avatar_contradiction_clause, [], [f8221])).
fof(f8221, plain, ($false | (~ spl10_61 | ~ spl10_121 | spl10_186)), inference(subsumption_resolution, [], [f8213, f1151])).
fof(f8213, plain, (~ (e0 = op(e0, e0)) | (~ spl10_61 | spl10_186)), inference(backward_demodulation, [], [f1523, f899])).
fof(f8199, plain, (~ spl10_56 | ~ spl10_81), inference(avatar_split_clause, [], [f8192, f981, f876])).
fof(f8192, plain, (~ (e0 = op(e2, e3)) | ~ spl10_81), inference(backward_demodulation, [], [f388, f983])).
fof(f8191, plain, (~ spl10_64 | ~ spl10_89), inference(avatar_split_clause, [], [f8186, f1014, f909])).
fof(f8186, plain, (~ (e3 = op(e2, e2)) | ~ spl10_89), inference(backward_demodulation, [], [f378, f1016])).
fof(f8182, plain, (spl10_3 | ~ spl10_95 | ~ spl10_230), inference(avatar_contradiction_clause, [], [f8181])).
fof(f8181, plain, ($false | (spl10_3 | ~ spl10_95 | ~ spl10_230)), inference(subsumption_resolution, [], [f8170, f654])).
fof(f654, plain, (~ (e2 = op(e4, e4)) | spl10_3), inference(avatar_component_clause, [], [f653])).
fof(f8170, plain, ((e2 = op(e4, e4)) | (~ spl10_95 | ~ spl10_230)), inference(backward_demodulation, [], [f1746, f1041])).
fof(f1746, plain, ((e2 = op(op(e1, e1), op(e1, e1))) | ~ spl10_230), inference(avatar_component_clause, [], [f1745])).
fof(f8104, plain, (~ spl10_57 | ~ spl10_4 | ~ spl10_8 | spl10_195), inference(avatar_split_clause, [], [f8093, f1561, f674, f657, f880])).
fof(f1561, plain, (spl10_195 <=> (e1 = op(op(e4, op(e4, e4)), op(e4, e4)))), introduced(avatar_definition, [new_symbols(naming, [spl10_195])])).
fof(f8093, plain, (~ (e1 = op(e2, e3)) | (~ spl10_4 | ~ spl10_8 | spl10_195)), inference(backward_demodulation, [], [f7859, f676])).
fof(f7859, plain, (~ (e1 = op(op(e4, e3), e3)) | (~ spl10_4 | spl10_195)), inference(forward_demodulation, [], [f1563, f659])).
fof(f1563, plain, (~ (e1 = op(op(e4, op(e4, e4)), op(e4, e4))) | spl10_195), inference(avatar_component_clause, [], [f1561])).
fof(f8103, plain, (~ spl10_56 | ~ spl10_4 | ~ spl10_8 | spl10_219), inference(avatar_split_clause, [], [f8094, f1681, f674, f657, f876])).
fof(f8094, plain, (~ (e0 = op(e2, e3)) | (~ spl10_4 | ~ spl10_8 | spl10_219)), inference(backward_demodulation, [], [f8000, f676])).
fof(f8000, plain, (~ (e0 = op(op(e4, e3), e3)) | (~ spl10_4 | spl10_219)), inference(forward_demodulation, [], [f1683, f659])).
fof(f8030, plain, (~ spl10_51 | ~ spl10_56), inference(avatar_split_clause, [], [f8020, f876, f855])).
fof(f8020, plain, (~ (e0 = op(e2, e4)) | ~ spl10_56), inference(backward_demodulation, [], [f433, f878])).
fof(f8015, plain, (~ spl10_18 | ~ spl10_12 | spl10_173), inference(avatar_split_clause, [], [f8014, f1449, f691, f716])).
fof(f8014, plain, (~ (e2 = op(e4, e1)) | (~ spl10_12 | spl10_173)), inference(forward_demodulation, [], [f1451, f693])).
fof(f7989, plain, (~ spl10_191 | ~ spl10_213), inference(avatar_contradiction_clause, [], [f7988])).
fof(f7988, plain, ($false | (~ spl10_191 | ~ spl10_213)), inference(subsumption_resolution, [], [f7987, f460])).
fof(f7987, plain, ((e1 = e4) | (~ spl10_191 | ~ spl10_213)), inference(forward_demodulation, [], [f1652, f1544])).
fof(f1544, plain, ((e1 = op(op(e3, op(e3, e3)), op(e3, e3))) | ~ spl10_191), inference(avatar_component_clause, [], [f1543])).
fof(f1652, plain, ((e4 = op(op(e3, op(e3, e3)), op(e3, e3))) | ~ spl10_213), inference(avatar_component_clause, [], [f1651])).
fof(f7978, plain, (~ spl10_32 | ~ spl10_64 | spl10_214), inference(avatar_split_clause, [], [f7496, f1657, f909, f775])).
fof(f7496, plain, (~ (e1 = op(e3, e3)) | (~ spl10_64 | spl10_214)), inference(backward_demodulation, [], [f1659, f911])).
fof(f7974, plain, (~ spl10_20 | ~ spl10_25), inference(avatar_split_clause, [], [f7973, f745, f724])).
fof(f7973, plain, (~ (e4 = op(e4, e1)) | ~ spl10_25), inference(forward_demodulation, [], [f444, f747])).
fof(f7972, plain, (~ spl10_20 | ~ spl10_70), inference(avatar_split_clause, [], [f7971, f934, f724])).
fof(f7971, plain, (~ (e4 = op(e4, e1)) | ~ spl10_70), inference(forward_demodulation, [], [f372, f936])).
fof(f7944, plain, (~ spl10_11 | ~ spl10_33 | ~ spl10_40 | spl10_217), inference(avatar_split_clause, [], [f7939, f1671, f808, f779, f687])).
fof(f7939, plain, (~ (e0 = op(e4, e2)) | (~ spl10_33 | ~ spl10_40 | spl10_217)), inference(backward_demodulation, [], [f7842, f810])).
fof(f7842, plain, (~ (e0 = op(op(e3, e2), e2)) | (~ spl10_33 | spl10_217)), inference(forward_demodulation, [], [f1673, f781])).
fof(f7943, plain, (~ spl10_40 | ~ spl10_51 | ~ spl10_170), inference(avatar_contradiction_clause, [], [f7942])).
fof(f7942, plain, ($false | (~ spl10_40 | ~ spl10_51 | ~ spl10_170)), inference(subsumption_resolution, [], [f7941, f456])).
fof(f7941, plain, ((e0 = e3) | (~ spl10_40 | ~ spl10_51 | ~ spl10_170)), inference(forward_demodulation, [], [f7937, f857])).
fof(f7937, plain, ((e3 = op(e2, e4)) | (~ spl10_40 | ~ spl10_170)), inference(backward_demodulation, [], [f1435, f810])).
fof(f7914, plain, (~ spl10_52 | ~ spl10_57), inference(avatar_split_clause, [], [f7906, f880, f859])).
fof(f7906, plain, (~ (e1 = op(e2, e4)) | ~ spl10_57), inference(backward_demodulation, [], [f433, f882])).
fof(f7904, plain, (~ spl10_52 | ~ spl10_70 | spl10_177), inference(avatar_split_clause, [], [f7901, f1470, f934, f859])).
fof(f7901, plain, (~ (e1 = op(e2, e4)) | (~ spl10_70 | spl10_177)), inference(backward_demodulation, [], [f1472, f936])).
fof(f7903, plain, (~ spl10_60 | ~ spl10_70), inference(avatar_split_clause, [], [f7898, f934, f892])).
fof(f7898, plain, (~ (e4 = op(e2, e3)) | ~ spl10_70), inference(backward_demodulation, [], [f429, f936])).
fof(f7902, plain, (~ spl10_45 | ~ spl10_70), inference(avatar_split_clause, [], [f7897, f934, f829])).
fof(f7897, plain, (~ (e4 = op(e3, e1)) | ~ spl10_70), inference(backward_demodulation, [], [f371, f936])).
fof(f7892, plain, (spl10_27 | ~ spl10_85 | ~ spl10_138), inference(avatar_split_clause, [], [f7887, f1273, f997, f754])).
fof(f7887, plain, ((e1 = op(e3, e4)) | (~ spl10_85 | ~ spl10_138)), inference(backward_demodulation, [], [f1275, f999])).
fof(f7891, plain, (~ spl10_60 | ~ spl10_85), inference(avatar_split_clause, [], [f7885, f997, f892])).
fof(f7885, plain, (~ (e4 = op(e2, e3)) | ~ spl10_85), inference(backward_demodulation, [], [f388, f999])).
fof(f7883, plain, (~ spl10_81 | ~ spl10_86), inference(avatar_split_clause, [], [f7875, f1002, f981])).
fof(f7875, plain, (~ (e0 = op(e1, e3)) | ~ spl10_86), inference(backward_demodulation, [], [f421, f1004])).
fof(f7882, plain, (~ spl10_11 | ~ spl10_86), inference(avatar_split_clause, [], [f7874, f1002, f687])).
fof(f7874, plain, (~ (e0 = op(e4, e2)) | ~ spl10_86), inference(backward_demodulation, [], [f380, f1004])).
fof(f7870, plain, (~ spl10_90 | ~ spl10_78 | spl10_133), inference(avatar_split_clause, [], [f7869, f1250, f968, f1018])).
fof(f7869, plain, (~ (e4 = op(e1, e2)) | (~ spl10_78 | spl10_133)), inference(forward_demodulation, [], [f1252, f970])).
fof(f7858, plain, (spl10_85 | ~ spl10_94 | ~ spl10_201), inference(avatar_split_clause, [], [f7857, f1588, f1035, f997])).
fof(f7857, plain, ((e4 = op(e1, e3)) | (~ spl10_94 | ~ spl10_201)), inference(forward_demodulation, [], [f1589, f1037])).
fof(f7852, plain, (~ spl10_7 | ~ spl10_4 | spl10_205), inference(avatar_split_clause, [], [f7851, f1607, f657, f670])).
fof(f7851, plain, (~ (e1 = op(e4, e3)) | (~ spl10_4 | spl10_205)), inference(forward_demodulation, [], [f1609, f659])).
fof(f7837, plain, (~ spl10_81 | ~ spl10_94 | spl10_237), inference(avatar_split_clause, [], [f7836, f1787, f1035, f981])).
fof(f7836, plain, (~ (e0 = op(e1, e3)) | (~ spl10_94 | spl10_237)), inference(forward_demodulation, [], [f1789, f1037])).
fof(f7829, plain, (~ spl10_57 | ~ spl10_64 | spl10_209), inference(avatar_split_clause, [], [f7495, f1627, f909, f880])).
fof(f7495, plain, (~ (e1 = op(e2, e3)) | (~ spl10_64 | spl10_209)), inference(backward_demodulation, [], [f1629, f911])).
fof(f1629, plain, (~ (e1 = op(e2, op(e2, e2))) | spl10_209), inference(avatar_component_clause, [], [f1627])).
fof(f7789, plain, (~ spl10_33 | ~ spl10_90 | ~ spl10_203 | spl10_213), inference(avatar_contradiction_clause, [], [f7788])).
fof(f7788, plain, ($false | (~ spl10_33 | ~ spl10_90 | ~ spl10_203 | spl10_213)), inference(subsumption_resolution, [], [f7785, f1020])).
fof(f7785, plain, (~ (e4 = op(e1, e2)) | (~ spl10_33 | ~ spl10_203 | spl10_213)), inference(backward_demodulation, [], [f7719, f781])).
fof(f7719, plain, (~ (e4 = op(e1, op(e3, e3))) | (~ spl10_203 | spl10_213)), inference(backward_demodulation, [], [f1653, f1598])).
fof(f7770, plain, (~ spl10_7 | ~ spl10_60 | ~ spl10_64 | spl10_187), inference(avatar_contradiction_clause, [], [f7769])).
fof(f7769, plain, ($false | (~ spl10_7 | ~ spl10_60 | ~ spl10_64 | spl10_187)), inference(subsumption_resolution, [], [f7765, f672])).
fof(f7765, plain, (~ (e1 = op(e4, e3)) | (~ spl10_60 | ~ spl10_64 | spl10_187)), inference(backward_demodulation, [], [f7626, f894])).
fof(f7626, plain, (~ (e1 = op(op(e2, e3), e3)) | (~ spl10_64 | spl10_187)), inference(forward_demodulation, [], [f1527, f911])).
fof(f7759, plain, (~ spl10_56 | ~ spl10_66), inference(avatar_split_clause, [], [f7751, f918, f876])).
fof(f7751, plain, (~ (e0 = op(e2, e3)) | ~ spl10_66), inference(backward_demodulation, [], [f429, f920])).
fof(f7715, plain, (~ spl10_56 | ~ spl10_64 | spl10_228), inference(avatar_split_clause, [], [f7714, f1727, f909, f876])).
fof(f7714, plain, (~ (e0 = op(e2, e3)) | (~ spl10_64 | spl10_228)), inference(forward_demodulation, [], [f1729, f911])).
fof(f7700, plain, (~ spl10_81 | ~ spl10_4 | ~ spl10_7 | spl10_219), inference(avatar_split_clause, [], [f7593, f1681, f670, f657, f981])).
fof(f7593, plain, (~ (e0 = op(e1, e3)) | (~ spl10_4 | ~ spl10_7 | spl10_219)), inference(backward_demodulation, [], [f6767, f672])).
fof(f6767, plain, (~ (e0 = op(op(e4, e3), e3)) | (~ spl10_4 | spl10_219)), inference(forward_demodulation, [], [f1683, f659])).
fof(f7681, plain, (~ spl10_28 | ~ spl10_35 | ~ spl10_52 | spl10_191), inference(avatar_contradiction_clause, [], [f7680])).
fof(f7680, plain, ($false | (~ spl10_28 | ~ spl10_35 | ~ spl10_52 | spl10_191)), inference(subsumption_resolution, [], [f7679, f861])).
fof(f7679, plain, (~ (e1 = op(e2, e4)) | (~ spl10_28 | ~ spl10_35 | spl10_191)), inference(forward_demodulation, [], [f7670, f760])).
fof(f7670, plain, (~ (e1 = op(op(e3, e4), e4)) | (~ spl10_35 | spl10_191)), inference(backward_demodulation, [], [f1545, f789])).
fof(f7639, plain, (spl10_1 | ~ spl10_105 | ~ spl10_135), inference(avatar_contradiction_clause, [], [f7638])).
fof(f7638, plain, ($false | (spl10_1 | ~ spl10_105 | ~ spl10_135)), inference(subsumption_resolution, [], [f7637, f646])).
fof(f646, plain, (~ (e0 = op(e4, e4)) | spl10_1), inference(avatar_component_clause, [], [f645])).
fof(f7637, plain, ((e0 = op(e4, e4)) | (~ spl10_105 | ~ spl10_135)), inference(forward_demodulation, [], [f1261, f1083])).
fof(f7636, plain, (~ spl10_105 | spl10_136), inference(avatar_contradiction_clause, [], [f7635])).
fof(f7635, plain, ($false | (~ spl10_105 | spl10_136)), inference(subsumption_resolution, [], [f7634, f1083])).
fof(f7634, plain, (~ (e4 = op(e0, e4)) | (~ spl10_105 | spl10_136)), inference(forward_demodulation, [], [f1266, f1083])).
fof(f7625, plain, (~ spl10_31 | ~ spl10_94 | spl10_198), inference(avatar_split_clause, [], [f7624, f1575, f1035, f771])).
fof(f7624, plain, (~ (e0 = op(e3, e3)) | (~ spl10_94 | spl10_198)), inference(forward_demodulation, [], [f1577, f1037])).
fof(f7613, plain, (~ spl10_83 | ~ spl10_94 | spl10_207), inference(avatar_contradiction_clause, [], [f7612])).
fof(f7612, plain, ($false | (~ spl10_83 | ~ spl10_94 | spl10_207)), inference(subsumption_resolution, [], [f7611, f991])).
fof(f7611, plain, (~ (e2 = op(e1, e3)) | (~ spl10_94 | spl10_207)), inference(forward_demodulation, [], [f1619, f1037])).
fof(f7610, plain, (~ spl10_60 | ~ spl10_83 | ~ spl10_94 | spl10_211), inference(avatar_split_clause, [], [f7609, f1639, f1035, f989, f892])).
fof(f7609, plain, (~ (e4 = op(e2, e3)) | (~ spl10_83 | ~ spl10_94 | spl10_211)), inference(forward_demodulation, [], [f7608, f991])).
fof(f7608, plain, (~ (e4 = op(op(e1, e3), e3)) | (~ spl10_94 | spl10_211)), inference(forward_demodulation, [], [f1641, f1037])).
fof(f7581, plain, (~ spl10_15 | ~ spl10_25), inference(avatar_split_clause, [], [f7578, f745, f703])).
fof(f7578, plain, (~ (e4 = op(e4, e2)) | ~ spl10_25), inference(backward_demodulation, [], [f445, f747])).
fof(f7565, plain, (~ spl10_31 | ~ spl10_49 | spl10_165), inference(avatar_contradiction_clause, [], [f7564])).
fof(f7564, plain, ($false | (~ spl10_31 | ~ spl10_49 | spl10_165)), inference(subsumption_resolution, [], [f7555, f848])).
fof(f7555, plain, (~ (e3 = op(e3, e0)) | (~ spl10_31 | spl10_165)), inference(backward_demodulation, [], [f1406, f773])).
fof(f7563, plain, (~ spl10_6 | ~ spl10_31), inference(avatar_split_clause, [], [f7549, f771, f666])).
fof(f7549, plain, (~ (e0 = op(e4, e3)) | ~ spl10_31), inference(backward_demodulation, [], [f393, f773])).
fof(f7534, plain, (~ spl10_31 | ~ spl10_49 | spl10_184), inference(avatar_split_clause, [], [f7532, f1507, f846, f771])).
fof(f7532, plain, (~ (e0 = op(e3, e3)) | (~ spl10_49 | spl10_184)), inference(backward_demodulation, [], [f1509, f848])).
fof(f7516, plain, (~ spl10_6 | ~ spl10_60 | ~ spl10_64 | spl10_215), inference(avatar_split_clause, [], [f7512, f1661, f909, f892, f666])).
fof(f7512, plain, (~ (e0 = op(e4, e3)) | (~ spl10_60 | ~ spl10_64 | spl10_215)), inference(backward_demodulation, [], [f7497, f894])).
fof(f7497, plain, (~ (e0 = op(op(e2, e3), e3)) | (~ spl10_64 | spl10_215)), inference(backward_demodulation, [], [f1663, f911])).
fof(f7502, plain, (~ spl10_31 | ~ spl10_64 | spl10_186), inference(avatar_split_clause, [], [f7492, f1521, f909, f771])).
fof(f7492, plain, (~ (e0 = op(e3, e3)) | (~ spl10_64 | spl10_186)), inference(backward_demodulation, [], [f1523, f911])).
fof(f7471, plain, (~ spl10_63 | ~ spl10_73), inference(avatar_split_clause, [], [f7464, f947, f905])).
fof(f7464, plain, (~ (e2 = op(e2, e2)) | ~ spl10_73), inference(backward_demodulation, [], [f425, f949])).
fof(f7459, plain, (spl10_52 | ~ spl10_90 | ~ spl10_144), inference(avatar_split_clause, [], [f7455, f1301, f1018, f859])).
fof(f7455, plain, ((e1 = op(e2, e4)) | (~ spl10_90 | ~ spl10_144)), inference(backward_demodulation, [], [f1303, f1020])).
fof(f7446, plain, (spl10_31 | ~ spl10_94 | ~ spl10_198), inference(avatar_split_clause, [], [f7442, f1575, f1035, f771])).
fof(f7442, plain, ((e0 = op(e3, e3)) | (~ spl10_94 | ~ spl10_198)), inference(backward_demodulation, [], [f1576, f1037])).
fof(f7437, plain, (~ spl10_87 | ~ spl10_97), inference(avatar_split_clause, [], [f7430, f1048, f1006])).
fof(f7430, plain, (~ (e1 = op(e1, e2)) | ~ spl10_97), inference(backward_demodulation, [], [f415, f1050])).
fof(f7425, plain, (~ spl10_55 | ~ spl10_105), inference(avatar_split_clause, [], [f7422, f1081, f871])).
fof(f7422, plain, (~ (e4 = op(e2, e4)) | ~ spl10_105), inference(backward_demodulation, [], [f395, f1083])).
fof(f7421, plain, (~ spl10_109 | spl10_142), inference(avatar_contradiction_clause, [], [f7420])).
fof(f7420, plain, ($false | (~ spl10_109 | spl10_142)), inference(subsumption_resolution, [], [f7417, f1100])).
fof(f7417, plain, (~ (e3 = op(e0, e3)) | (~ spl10_109 | spl10_142)), inference(backward_demodulation, [], [f1294, f1100])).
fof(f7418, plain, (~ spl10_59 | ~ spl10_109), inference(avatar_split_clause, [], [f7411, f1098, f888])).
fof(f7411, plain, (~ (e3 = op(e2, e3)) | ~ spl10_109), inference(backward_demodulation, [], [f385, f1100])).
fof(f7409, plain, (~ spl10_63 | ~ spl10_113), inference(avatar_split_clause, [], [f7401, f1115, f905])).
fof(f7401, plain, (~ (e2 = op(e2, e2)) | ~ spl10_113), inference(backward_demodulation, [], [f375, f1117])).
fof(f7398, plain, (~ spl10_67 | ~ spl10_117), inference(avatar_split_clause, [], [f7388, f1132, f922])).
fof(f7388, plain, (~ (e1 = op(e2, e1)) | ~ spl10_117), inference(backward_demodulation, [], [f365, f1134])).
fof(f7379, plain, (~ spl10_111 | ~ spl10_121), inference(avatar_split_clause, [], [f7365, f1149, f1107])).
fof(f7365, plain, (~ (e0 = op(e0, e2)) | ~ spl10_121), inference(backward_demodulation, [], [f405, f1151])).
fof(f7378, plain, (~ spl10_71 | ~ spl10_121), inference(avatar_split_clause, [], [f7361, f1149, f939])).
fof(f7361, plain, (~ (e0 = op(e2, e0)) | ~ spl10_121), inference(backward_demodulation, [], [f355, f1151])).
fof(f7335, plain, (spl10_37 | ~ spl10_83 | ~ spl10_138), inference(avatar_split_clause, [], [f7334, f1273, f989, f796])).
fof(f7334, plain, ((e1 = op(e3, e2)) | (~ spl10_83 | ~ spl10_138)), inference(forward_demodulation, [], [f1275, f991])).
fof(f7316, plain, (~ spl10_100 | ~ spl10_76 | spl10_133), inference(avatar_split_clause, [], [f7171, f1250, f960, f1060])).
fof(f7171, plain, (~ (e4 = op(e1, e0)) | (~ spl10_76 | spl10_133)), inference(backward_demodulation, [], [f1252, f962])).
fof(f7311, plain, (~ spl10_89 | ~ spl10_83 | spl10_139), inference(avatar_split_clause, [], [f7159, f1278, f989, f1014])).
fof(f7159, plain, (~ (e3 = op(e1, e2)) | (~ spl10_83 | spl10_139)), inference(backward_demodulation, [], [f1280, f991])).
fof(f7309, plain, (~ spl10_4 | ~ spl10_6 | ~ spl10_107 | spl10_195), inference(avatar_contradiction_clause, [], [f7308])).
fof(f7308, plain, ($false | (~ spl10_4 | ~ spl10_6 | ~ spl10_107 | spl10_195)), inference(subsumption_resolution, [], [f7301, f1092])).
fof(f7301, plain, (~ (e1 = op(e0, e3)) | (~ spl10_4 | ~ spl10_6 | spl10_195)), inference(backward_demodulation, [], [f6080, f668])).
fof(f6080, plain, (~ (e1 = op(op(e4, e3), e3)) | (~ spl10_4 | spl10_195)), inference(forward_demodulation, [], [f1563, f659])).
fof(f7306, plain, (~ spl10_6 | ~ spl10_48 | ~ spl10_166), inference(avatar_contradiction_clause, [], [f7305])).
fof(f7305, plain, ($false | (~ spl10_6 | ~ spl10_48 | ~ spl10_166)), inference(subsumption_resolution, [], [f7304, f462])).
fof(f462, plain, ~ (e2 = e4), inference(cnf_transformation, [], [f5])).
fof(f7304, plain, ((e2 = e4) | (~ spl10_6 | ~ spl10_48 | ~ spl10_166)), inference(forward_demodulation, [], [f7299, f844])).
fof(f7299, plain, ((e4 = op(e3, e0)) | (~ spl10_6 | ~ spl10_166)), inference(backward_demodulation, [], [f1411, f668])).
fof(f7282, plain, (~ spl10_27 | ~ spl10_35 | ~ spl10_76 | spl10_217), inference(avatar_contradiction_clause, [], [f7281])).
fof(f7281, plain, ($false | (~ spl10_27 | ~ spl10_35 | ~ spl10_76 | spl10_217)), inference(subsumption_resolution, [], [f7275, f962])).
fof(f7275, plain, (~ (e0 = op(e1, e4)) | (~ spl10_27 | ~ spl10_35 | spl10_217)), inference(backward_demodulation, [], [f7259, f756])).
fof(f7259, plain, (~ (e0 = op(op(e3, e4), e4)) | (~ spl10_35 | spl10_217)), inference(backward_demodulation, [], [f1673, f789])).
fof(f7263, plain, (~ spl10_26 | ~ spl10_35 | spl10_224), inference(avatar_split_clause, [], [f7260, f1705, f787, f750])).
fof(f7260, plain, (~ (e0 = op(e3, e4)) | (~ spl10_35 | spl10_224)), inference(backward_demodulation, [], [f1707, f789])).
fof(f1707, plain, (~ (e0 = op(e3, op(e3, e3))) | spl10_224), inference(avatar_component_clause, [], [f1705])).
fof(f7219, plain, (~ spl10_55 | spl10_162), inference(avatar_contradiction_clause, [], [f7218])).
fof(f7218, plain, ($false | (~ spl10_55 | spl10_162)), inference(subsumption_resolution, [], [f7213, f873])).
fof(f7213, plain, (~ (e4 = op(e2, e4)) | (~ spl10_55 | spl10_162)), inference(backward_demodulation, [], [f1388, f873])).
fof(f7217, plain, (spl10_3 | ~ spl10_55 | ~ spl10_159), inference(avatar_contradiction_clause, [], [f7216])).
fof(f7216, plain, ($false | (spl10_3 | ~ spl10_55 | ~ spl10_159)), inference(subsumption_resolution, [], [f7212, f654])).
fof(f7212, plain, ((e2 = op(e4, e4)) | (~ spl10_55 | ~ spl10_159)), inference(backward_demodulation, [], [f1375, f873])).
fof(f7199, plain, (~ spl10_52 | ~ spl10_67), inference(avatar_split_clause, [], [f7193, f922, f859])).
fof(f7193, plain, (~ (e1 = op(e2, e4)) | ~ spl10_67), inference(backward_demodulation, [], [f430, f924])).
fof(f7190, plain, (~ spl10_71 | spl10_183), inference(avatar_contradiction_clause, [], [f7189])).
fof(f7189, plain, ($false | (~ spl10_71 | spl10_183)), inference(subsumption_resolution, [], [f7185, f941])).
fof(f7185, plain, (~ (e0 = op(e2, e0)) | (~ spl10_71 | spl10_183)), inference(backward_demodulation, [], [f1504, f941])).
fof(f7186, plain, (~ spl10_56 | ~ spl10_71), inference(avatar_split_clause, [], [f7178, f939, f876])).
fof(f7178, plain, (~ (e0 = op(e2, e3)) | ~ spl10_71), inference(backward_demodulation, [], [f426, f941])).
fof(f7150, plain, (~ spl10_37 | ~ spl10_87), inference(avatar_split_clause, [], [f7141, f1006, f796])).
fof(f7141, plain, (~ (e1 = op(e3, e2)) | ~ spl10_87), inference(backward_demodulation, [], [f379, f1008])).
fof(f7138, plain, (~ spl10_95 | ~ spl10_103 | spl10_199 | ~ spl10_237), inference(avatar_contradiction_clause, [], [f7137])).
fof(f7137, plain, ($false | (~ spl10_95 | ~ spl10_103 | spl10_199 | ~ spl10_237)), inference(subsumption_resolution, [], [f7131, f1075])).
fof(f7131, plain, (~ (e2 = op(e0, e4)) | (~ spl10_95 | spl10_199 | ~ spl10_237)), inference(backward_demodulation, [], [f7005, f1041])).
fof(f7005, plain, (~ (e2 = op(e0, op(e1, e1))) | (spl10_199 | ~ spl10_237)), inference(backward_demodulation, [], [f1581, f1788])).
fof(f7120, plain, (~ spl10_49 | ~ spl10_99), inference(avatar_split_clause, [], [f7111, f1056, f846])).
fof(f7111, plain, (~ (e3 = op(e3, e0)) | ~ spl10_99), inference(backward_demodulation, [], [f359, f1058])).
fof(f7096, plain, (~ spl10_7 | ~ spl10_107), inference(avatar_split_clause, [], [f7090, f1090, f670])).
fof(f7090, plain, (~ (e1 = op(e4, e3)) | ~ spl10_107), inference(backward_demodulation, [], [f387, f1092])).
fof(f7069, plain, (~ spl10_69 | ~ spl10_119), inference(avatar_split_clause, [], [f7061, f1140, f930])).
fof(f7061, plain, (~ (e3 = op(e2, e1)) | ~ spl10_119), inference(backward_demodulation, [], [f365, f1142])).
fof(f7013, plain, (~ spl10_52 | ~ spl10_65 | spl10_209), inference(avatar_split_clause, [], [f7012, f1627, f913, f859])).
fof(f7012, plain, (~ (e1 = op(e2, e4)) | (~ spl10_65 | spl10_209)), inference(forward_demodulation, [], [f1629, f915])).
fof(f6996, plain, (spl10_90 | ~ spl10_18 | ~ spl10_176), inference(avatar_split_clause, [], [f6892, f1465, f716, f1018])).
fof(f6892, plain, ((e4 = op(e1, e2)) | (~ spl10_18 | ~ spl10_176)), inference(forward_demodulation, [], [f1467, f718])).
fof(f6988, plain, (~ spl10_40 | ~ spl10_65), inference(avatar_split_clause, [], [f6987, f913, f808])).
fof(f6987, plain, (~ (e4 = op(e3, e2)) | ~ spl10_65), inference(forward_demodulation, [], [f381, f915])).
fof(f6973, plain, (~ spl10_33 | ~ spl10_37 | spl10_203), inference(avatar_contradiction_clause, [], [f6972])).
fof(f6972, plain, ($false | (~ spl10_33 | ~ spl10_37 | spl10_203)), inference(subsumption_resolution, [], [f6959, f798])).
fof(f6959, plain, (~ (e1 = op(e3, e2)) | (~ spl10_33 | spl10_203)), inference(backward_demodulation, [], [f1599, f781])).
fof(f6905, plain, (spl10_91 | ~ spl10_117 | ~ spl10_152), inference(avatar_split_clause, [], [f6904, f1339, f1132, f1023])).
fof(f6904, plain, ((e0 = op(e1, e1)) | (~ spl10_117 | ~ spl10_152)), inference(forward_demodulation, [], [f1341, f1134])).
fof(f6897, plain, (~ spl10_83 | ~ spl10_69 | spl10_174), inference(avatar_split_clause, [], [f6896, f1457, f930, f989])).
fof(f6896, plain, (~ (e2 = op(e1, e3)) | (~ spl10_69 | spl10_174)), inference(forward_demodulation, [], [f1458, f932])).
fof(f1458, plain, (~ (e2 = op(e1, op(e2, e1))) | spl10_174), inference(avatar_component_clause, [], [f1457])).
fof(f6895, plain, (~ spl10_18 | ~ spl10_89 | ~ spl10_176), inference(avatar_contradiction_clause, [], [f6894])).
fof(f6894, plain, ($false | (~ spl10_18 | ~ spl10_89 | ~ spl10_176)), inference(subsumption_resolution, [], [f6893, f463])).
fof(f6893, plain, ((e3 = e4) | (~ spl10_18 | ~ spl10_89 | ~ spl10_176)), inference(forward_demodulation, [], [f6892, f1016])).
fof(f6851, plain, (~ spl10_11 | spl10_23 | ~ spl10_173), inference(avatar_contradiction_clause, [], [f6850])).
fof(f6850, plain, ($false | (~ spl10_11 | spl10_23 | ~ spl10_173)), inference(subsumption_resolution, [], [f6847, f738])).
fof(f738, plain, (~ (e2 = op(e4, e0)) | spl10_23), inference(avatar_component_clause, [], [f737])).
fof(f6847, plain, ((e2 = op(e4, e0)) | (~ spl10_11 | ~ spl10_173)), inference(backward_demodulation, [], [f1450, f689])).
fof(f6831, plain, (~ spl10_32 | ~ spl10_35), inference(avatar_contradiction_clause, [], [f6830])).
fof(f6830, plain, ($false | (~ spl10_32 | ~ spl10_35)), inference(subsumption_resolution, [], [f6829, f460])).
fof(f6829, plain, ((e1 = e4) | (~ spl10_32 | ~ spl10_35)), inference(forward_demodulation, [], [f789, f777])).
fof(f6816, plain, (~ spl10_65 | ~ spl10_76 | ~ spl10_187 | ~ spl10_209), inference(avatar_contradiction_clause, [], [f6815])).
fof(f6815, plain, ($false | (~ spl10_65 | ~ spl10_76 | ~ spl10_187 | ~ spl10_209)), inference(subsumption_resolution, [], [f6814, f454])).
fof(f454, plain, ~ (e0 = e1), inference(cnf_transformation, [], [f5])).
fof(f6814, plain, ((e0 = e1) | (~ spl10_65 | ~ spl10_76 | ~ spl10_187 | ~ spl10_209)), inference(forward_demodulation, [], [f6810, f962])).
fof(f6810, plain, ((e1 = op(e1, e4)) | (~ spl10_65 | ~ spl10_187 | ~ spl10_209)), inference(backward_demodulation, [], [f6773, f915])).
fof(f6773, plain, ((e1 = op(e1, op(e2, e2))) | (~ spl10_187 | ~ spl10_209)), inference(forward_demodulation, [], [f1526, f1628])).
fof(f1526, plain, ((e1 = op(op(e2, op(e2, e2)), op(e2, e2))) | ~ spl10_187), inference(avatar_component_clause, [], [f1525])).
fof(f6766, plain, (~ spl10_32 | ~ spl10_41 | spl10_224), inference(avatar_contradiction_clause, [], [f6765])).
fof(f6765, plain, ($false | (~ spl10_32 | ~ spl10_41 | spl10_224)), inference(subsumption_resolution, [], [f6764, f815])).
fof(f6764, plain, (~ (e0 = op(e3, e1)) | (~ spl10_32 | spl10_224)), inference(forward_demodulation, [], [f1707, f777])).
fof(f6734, plain, (~ spl10_41 | ~ spl10_97 | ~ spl10_175), inference(avatar_contradiction_clause, [], [f6733])).
fof(f6733, plain, ($false | (~ spl10_41 | ~ spl10_97 | ~ spl10_175)), inference(subsumption_resolution, [], [f6732, f459])).
fof(f6732, plain, ((e1 = e3) | (~ spl10_41 | ~ spl10_97 | ~ spl10_175)), inference(forward_demodulation, [], [f6731, f1050])).
fof(f6731, plain, ((e3 = op(e1, e0)) | (~ spl10_41 | ~ spl10_175)), inference(forward_demodulation, [], [f1463, f815])).
fof(f6720, plain, (~ spl10_93 | ~ spl10_32 | spl10_232), inference(avatar_split_clause, [], [f6719, f1755, f775, f1031])).
fof(f6719, plain, (~ (e2 = op(e1, e1)) | (~ spl10_32 | spl10_232)), inference(forward_demodulation, [], [f1757, f777])).
fof(f6684, plain, (~ spl10_32 | ~ spl10_33), inference(avatar_contradiction_clause, [], [f6683])).
fof(f6683, plain, ($false | (~ spl10_32 | ~ spl10_33)), inference(subsumption_resolution, [], [f6682, f458])).
fof(f6682, plain, ((e1 = e2) | (~ spl10_32 | ~ spl10_33)), inference(backward_demodulation, [], [f781, f777])).
fof(f6677, plain, (~ spl10_12 | ~ spl10_37), inference(avatar_split_clause, [], [f6669, f796, f691])).
fof(f6669, plain, (~ (e1 = op(e4, e2)) | ~ spl10_37), inference(backward_demodulation, [], [f383, f798])).
fof(f6668, plain, (~ spl10_41 | ~ spl10_45), inference(avatar_contradiction_clause, [], [f6667])).
fof(f6667, plain, ($false | (~ spl10_41 | ~ spl10_45)), inference(subsumption_resolution, [], [f6666, f457])).
fof(f6666, plain, ((e0 = e4) | (~ spl10_41 | ~ spl10_45)), inference(backward_demodulation, [], [f831, f815])).
fof(f6640, plain, (~ spl10_51 | ~ spl10_61), inference(avatar_split_clause, [], [f6626, f897, f855])).
fof(f6626, plain, (~ (e0 = op(e2, e4)) | ~ spl10_61), inference(backward_demodulation, [], [f432, f899])).
fof(f6621, plain, (~ spl10_51 | ~ spl10_76), inference(avatar_split_clause, [], [f6613, f960, f855])).
fof(f6613, plain, (~ (e0 = op(e2, e4)) | ~ spl10_76), inference(backward_demodulation, [], [f398, f962])).
fof(f6602, plain, (~ spl10_89 | ~ spl10_90), inference(avatar_contradiction_clause, [], [f6601])).
fof(f6601, plain, ($false | (~ spl10_89 | ~ spl10_90)), inference(subsumption_resolution, [], [f6600, f463])).
fof(f6600, plain, ((e3 = e4) | (~ spl10_89 | ~ spl10_90)), inference(forward_demodulation, [], [f1020, f1016])).
fof(f6597, plain, (~ spl10_89 | ~ spl10_93 | spl10_200), inference(avatar_contradiction_clause, [], [f6596])).
fof(f6596, plain, ($false | (~ spl10_89 | ~ spl10_93 | spl10_200)), inference(subsumption_resolution, [], [f6587, f1016])).
fof(f6587, plain, (~ (e3 = op(e1, e2)) | (~ spl10_93 | spl10_200)), inference(backward_demodulation, [], [f1585, f1033])).
fof(f6573, plain, (spl10_83 | ~ spl10_69 | ~ spl10_174), inference(avatar_split_clause, [], [f6572, f1457, f930, f989])).
fof(f6572, plain, ((e2 = op(e1, e3)) | (~ spl10_69 | ~ spl10_174)), inference(forward_demodulation, [], [f1459, f932])).
fof(f6563, plain, (~ spl10_65 | ~ spl10_33 | spl10_244), inference(avatar_split_clause, [], [f6562, f1853, f779, f913])).
fof(f6562, plain, (~ (e4 = op(e2, e2)) | (~ spl10_33 | spl10_244)), inference(forward_demodulation, [], [f1855, f781])).
fof(f6556, plain, (spl10_79 | ~ spl10_45 | ~ spl10_175), inference(avatar_split_clause, [], [f6463, f1461, f829, f972])).
fof(f6463, plain, ((e3 = op(e1, e4)) | (~ spl10_45 | ~ spl10_175)), inference(backward_demodulation, [], [f1463, f831])).
fof(f6547, plain, (~ spl10_58 | ~ spl10_73), inference(avatar_split_clause, [], [f6546, f947, f884])).
fof(f6546, plain, (~ (e2 = op(e2, e3)) | ~ spl10_73), inference(forward_demodulation, [], [f426, f949])).
fof(f6526, plain, (~ spl10_38 | ~ spl10_113), inference(avatar_split_clause, [], [f6525, f1115, f800])).
fof(f6525, plain, (~ (e2 = op(e3, e2)) | ~ spl10_113), inference(forward_demodulation, [], [f376, f1117])).
fof(f6516, plain, (~ spl10_30 | ~ spl10_105), inference(avatar_split_clause, [], [f6515, f1081, f766])).
fof(f6515, plain, (~ (e4 = op(e3, e4)) | ~ spl10_105), inference(forward_demodulation, [], [f396, f1083])).
fof(f6514, plain, (~ spl10_27 | ~ spl10_45 | spl10_178), inference(avatar_split_clause, [], [f6464, f1475, f829, f754])).
fof(f6464, plain, (~ (e1 = op(e3, e4)) | (~ spl10_45 | spl10_178)), inference(backward_demodulation, [], [f1477, f831])).
fof(f6501, plain, (~ spl10_6 | ~ spl10_49 | ~ spl10_166), inference(avatar_contradiction_clause, [], [f6500])).
fof(f6500, plain, ($false | (~ spl10_6 | ~ spl10_49 | ~ spl10_166)), inference(subsumption_resolution, [], [f6499, f463])).
fof(f6499, plain, ((e3 = e4) | (~ spl10_6 | ~ spl10_49 | ~ spl10_166)), inference(forward_demodulation, [], [f6494, f848])).
fof(f6494, plain, ((e4 = op(e3, e0)) | (~ spl10_6 | ~ spl10_166)), inference(backward_demodulation, [], [f1411, f668])).
fof(f6468, plain, (~ spl10_45 | ~ spl10_78 | ~ spl10_175), inference(avatar_contradiction_clause, [], [f6467])).
fof(f6467, plain, ($false | (~ spl10_45 | ~ spl10_78 | ~ spl10_175)), inference(subsumption_resolution, [], [f6466, f461])).
fof(f6466, plain, ((e2 = e3) | (~ spl10_45 | ~ spl10_78 | ~ spl10_175)), inference(forward_demodulation, [], [f6463, f970])).
fof(f6354, plain, (~ spl10_16 | ~ spl10_97 | ~ spl10_176), inference(avatar_contradiction_clause, [], [f6353])).
fof(f6353, plain, ($false | (~ spl10_16 | ~ spl10_97 | ~ spl10_176)), inference(subsumption_resolution, [], [f6352, f460])).
fof(f6352, plain, ((e1 = e4) | (~ spl10_16 | ~ spl10_97 | ~ spl10_176)), inference(forward_demodulation, [], [f6349, f1050])).
fof(f6349, plain, ((e4 = op(e1, e0)) | (~ spl10_16 | ~ spl10_176)), inference(backward_demodulation, [], [f1467, f710])).
fof(f6322, plain, (~ spl10_44 | ~ spl10_49), inference(avatar_split_clause, [], [f6317, f846, f825])).
fof(f6317, plain, (~ (e3 = op(e3, e1)) | ~ spl10_49), inference(backward_demodulation, [], [f434, f848])).
fof(f6306, plain, (~ spl10_68 | ~ spl10_73), inference(avatar_split_clause, [], [f6300, f947, f926])).
fof(f6300, plain, (~ (e2 = op(e2, e1)) | ~ spl10_73), inference(backward_demodulation, [], [f424, f949])).
fof(f6297, plain, (~ spl10_51 | ~ spl10_78 | ~ spl10_95 | spl10_231), inference(avatar_contradiction_clause, [], [f6296])).
fof(f6296, plain, ($false | (~ spl10_51 | ~ spl10_78 | ~ spl10_95 | spl10_231)), inference(subsumption_resolution, [], [f6292, f857])).
fof(f6292, plain, (~ (e0 = op(e2, e4)) | (~ spl10_78 | ~ spl10_95 | spl10_231)), inference(backward_demodulation, [], [f6258, f970])).
fof(f6258, plain, (~ (e0 = op(op(e1, e4), e4)) | (~ spl10_95 | spl10_231)), inference(backward_demodulation, [], [f1751, f1041])).
fof(f6268, plain, (~ spl10_84 | ~ spl10_89), inference(avatar_split_clause, [], [f6263, f1014, f993])).
fof(f6263, plain, (~ (e3 = op(e1, e3)) | ~ spl10_89), inference(backward_demodulation, [], [f421, f1016])).
fof(f6222, plain, (~ spl10_113 | spl10_148), inference(avatar_contradiction_clause, [], [f6221])).
fof(f6221, plain, ($false | (~ spl10_113 | spl10_148)), inference(subsumption_resolution, [], [f6218, f1117])).
fof(f6218, plain, (~ (e2 = op(e0, e2)) | (~ spl10_113 | spl10_148)), inference(backward_demodulation, [], [f1322, f1117])).
fof(f6219, plain, (~ spl10_88 | ~ spl10_113), inference(avatar_split_clause, [], [f6212, f1115, f1010])).
fof(f6212, plain, (~ (e2 = op(e1, e2)) | ~ spl10_113), inference(backward_demodulation, [], [f374, f1117])).
fof(f6160, plain, (~ spl10_117 | ~ spl10_122 | spl10_236), inference(avatar_split_clause, [], [f5145, f1781, f1153, f1132])).
fof(f5145, plain, (~ (e1 = op(e0, e1)) | (~ spl10_122 | spl10_236)), inference(backward_demodulation, [], [f1783, f1155])).
fof(f6159, plain, (~ spl10_117 | ~ spl10_122), inference(avatar_split_clause, [], [f6158, f1153, f1132])).
fof(f6158, plain, (~ (e1 = op(e0, e1)) | ~ spl10_122), inference(forward_demodulation, [], [f404, f1155])).
fof(f6150, plain, (~ spl10_104 | ~ spl10_65 | spl10_208 | ~ spl10_228), inference(avatar_split_clause, [], [f5992, f1727, f1623, f913, f1077])).
fof(f1623, plain, (spl10_208 <=> (e3 = op(op(e2, op(e2, e2)), op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl10_208])])).
fof(f5992, plain, (~ (e3 = op(e0, e4)) | (~ spl10_65 | spl10_208 | ~ spl10_228)), inference(backward_demodulation, [], [f5973, f915])).
fof(f5973, plain, (~ (e3 = op(e0, op(e2, e2))) | (spl10_208 | ~ spl10_228)), inference(forward_demodulation, [], [f1625, f1728])).
fof(f1625, plain, (~ (e3 = op(op(e2, op(e2, e2)), op(e2, e2))) | spl10_208), inference(avatar_component_clause, [], [f1623])).
fof(f6149, plain, (~ spl10_97 | ~ spl10_122), inference(avatar_split_clause, [], [f5131, f1153, f1048])).
fof(f5131, plain, (~ (e1 = op(e1, e0)) | ~ spl10_122), inference(backward_demodulation, [], [f354, f1155])).
fof(f6148, plain, (~ spl10_95 | ~ spl10_122 | spl10_245), inference(avatar_split_clause, [], [f5434, f1859, f1153, f1039])).
fof(f5434, plain, (~ (e4 = op(e1, e1)) | (~ spl10_122 | spl10_245)), inference(forward_demodulation, [], [f1861, f1155])).
fof(f6146, plain, (spl10_92 | ~ spl10_122 | ~ spl10_220), inference(avatar_split_clause, [], [f5439, f1687, f1153, f1027])).
fof(f5439, plain, ((e1 = op(e1, e1)) | (~ spl10_122 | ~ spl10_220)), inference(forward_demodulation, [], [f1688, f1155])).
fof(f6103, plain, (~ spl10_23 | ~ spl10_114 | ~ spl10_182), inference(avatar_contradiction_clause, [], [f6102])).
fof(f6102, plain, ($false | (~ spl10_23 | ~ spl10_114 | ~ spl10_182)), inference(subsumption_resolution, [], [f6101, f463])).
fof(f6101, plain, ((e3 = e4) | (~ spl10_23 | ~ spl10_114 | ~ spl10_182)), inference(forward_demodulation, [], [f6100, f1121])).
fof(f6100, plain, ((e4 = op(e0, e2)) | (~ spl10_23 | ~ spl10_182)), inference(forward_demodulation, [], [f1499, f739])).
fof(f6093, plain, (~ spl10_33 | ~ spl10_36 | spl10_111 | ~ spl10_217), inference(avatar_contradiction_clause, [], [f6092])).
fof(f6092, plain, ($false | (~ spl10_33 | ~ spl10_36 | spl10_111 | ~ spl10_217)), inference(subsumption_resolution, [], [f6091, f1108])).
fof(f1108, plain, (~ (e0 = op(e0, e2)) | spl10_111), inference(avatar_component_clause, [], [f1107])).
fof(f6091, plain, ((e0 = op(e0, e2)) | (~ spl10_33 | ~ spl10_36 | ~ spl10_217)), inference(forward_demodulation, [], [f6090, f794])).
fof(f6090, plain, ((e0 = op(op(e3, e2), e2)) | (~ spl10_33 | ~ spl10_217)), inference(forward_demodulation, [], [f1672, f781])).
fof(f1672, plain, ((e0 = op(op(e3, op(e3, e3)), op(e3, e3))) | ~ spl10_217), inference(avatar_component_clause, [], [f1671])).
fof(f6049, plain, (~ spl10_33 | ~ spl10_36 | spl10_224), inference(avatar_contradiction_clause, [], [f6048])).
fof(f6048, plain, ($false | (~ spl10_33 | ~ spl10_36 | spl10_224)), inference(subsumption_resolution, [], [f6036, f794])).
fof(f6036, plain, (~ (e0 = op(e3, e2)) | (~ spl10_33 | spl10_224)), inference(backward_demodulation, [], [f1707, f781])).
fof(f6005, plain, (~ spl10_1 | ~ spl10_51), inference(avatar_split_clause, [], [f5998, f855, f645])).
fof(f5998, plain, (~ (e0 = op(e4, e4)) | ~ spl10_51), inference(backward_demodulation, [], [f402, f857])).
fof(f5996, plain, (spl10_51 | ~ spl10_65 | ~ spl10_228), inference(avatar_split_clause, [], [f5989, f1727, f913, f855])).
fof(f5989, plain, ((e0 = op(e2, e4)) | (~ spl10_65 | ~ spl10_228)), inference(backward_demodulation, [], [f1728, f915])).
fof(f5939, plain, (~ spl10_61 | ~ spl10_75 | spl10_189), inference(avatar_contradiction_clause, [], [f5938])).
fof(f5938, plain, ($false | (~ spl10_61 | ~ spl10_75 | spl10_189)), inference(subsumption_resolution, [], [f5937, f957])).
fof(f5937, plain, (~ (e4 = op(e2, e0)) | (~ spl10_61 | spl10_189)), inference(forward_demodulation, [], [f1536, f899])).
fof(f5931, plain, (~ spl10_1 | ~ spl10_122 | spl10_218), inference(avatar_contradiction_clause, [], [f5930])).
fof(f5930, plain, ($false | (~ spl10_1 | ~ spl10_122 | spl10_218)), inference(subsumption_resolution, [], [f5929, f1155])).
fof(f5929, plain, (~ (op(e0, e0) = e1) | (~ spl10_1 | spl10_218)), inference(forward_demodulation, [], [f1679, f647])).
fof(f5922, plain, (~ spl10_31 | ~ spl10_48 | ~ spl10_75 | spl10_213), inference(avatar_contradiction_clause, [], [f5921])).
fof(f5921, plain, ($false | (~ spl10_31 | ~ spl10_48 | ~ spl10_75 | spl10_213)), inference(subsumption_resolution, [], [f5920, f957])).
fof(f5920, plain, (~ (e4 = op(e2, e0)) | (~ spl10_31 | ~ spl10_48 | spl10_213)), inference(forward_demodulation, [], [f5919, f844])).
fof(f5919, plain, (~ (e4 = op(op(e3, e0), e0)) | (~ spl10_31 | spl10_213)), inference(forward_demodulation, [], [f1653, f773])).
fof(f5914, plain, (~ spl10_1 | ~ spl10_24 | ~ spl10_48 | spl10_204), inference(avatar_contradiction_clause, [], [f5913])).
fof(f5913, plain, ($false | (~ spl10_1 | ~ spl10_24 | ~ spl10_48 | spl10_204)), inference(subsumption_resolution, [], [f5912, f844])).
fof(f5912, plain, (~ (e2 = op(e3, e0)) | (~ spl10_1 | ~ spl10_24 | spl10_204)), inference(forward_demodulation, [], [f5904, f743])).
fof(f5904, plain, (~ (e2 = op(op(e4, e0), e0)) | (~ spl10_1 | spl10_204)), inference(backward_demodulation, [], [f1605, f647])).
fof(f5911, plain, (~ spl10_1 | ~ spl10_24 | spl10_197), inference(avatar_contradiction_clause, [], [f5910])).
fof(f5910, plain, ($false | (~ spl10_1 | ~ spl10_24 | spl10_197)), inference(subsumption_resolution, [], [f5903, f743])).
fof(f5903, plain, (~ (e3 = op(e4, e0)) | (~ spl10_1 | spl10_197)), inference(backward_demodulation, [], [f1572, f647])).
fof(f1572, plain, (~ (e3 = op(e4, op(e4, e4))) | spl10_197), inference(avatar_component_clause, [], [f1570])).
fof(f5876, plain, (~ spl10_31 | ~ spl10_48 | spl10_192), inference(avatar_contradiction_clause, [], [f5875])).
fof(f5875, plain, ($false | (~ spl10_31 | ~ spl10_48 | spl10_192)), inference(subsumption_resolution, [], [f5868, f844])).
fof(f5868, plain, (~ (e2 = op(e3, e0)) | (~ spl10_31 | spl10_192)), inference(backward_demodulation, [], [f1549, f773])).
fof(f5874, plain, (~ spl10_31 | ~ spl10_48 | ~ spl10_165), inference(avatar_contradiction_clause, [], [f5873])).
fof(f5873, plain, ($false | (~ spl10_31 | ~ spl10_48 | ~ spl10_165)), inference(subsumption_resolution, [], [f5872, f461])).
fof(f5872, plain, ((e2 = e3) | (~ spl10_31 | ~ spl10_48 | ~ spl10_165)), inference(forward_demodulation, [], [f5867, f844])).
fof(f5867, plain, ((e3 = op(e3, e0)) | (~ spl10_31 | ~ spl10_165)), inference(backward_demodulation, [], [f1407, f773])).
fof(f5849, plain, (~ spl10_61 | ~ spl10_122 | spl10_214), inference(avatar_contradiction_clause, [], [f5848])).
fof(f5848, plain, ($false | (~ spl10_61 | ~ spl10_122 | spl10_214)), inference(subsumption_resolution, [], [f5840, f1155])).
fof(f5840, plain, (~ (op(e0, e0) = e1) | (~ spl10_61 | spl10_214)), inference(backward_demodulation, [], [f1659, f899])).
fof(f5814, plain, (~ spl10_203 | ~ spl10_165 | spl10_191), inference(avatar_split_clause, [], [f5813, f1543, f1405, f1597])).
fof(f5813, plain, (~ (e1 = op(e3, op(e3, e3))) | (~ spl10_165 | spl10_191)), inference(forward_demodulation, [], [f1545, f1407])).
fof(f5811, plain, (~ spl10_224 | ~ spl10_165 | spl10_217), inference(avatar_split_clause, [], [f5810, f1671, f1405, f1705])).
fof(f5810, plain, (~ (e0 = op(e3, op(e3, e3))) | (~ spl10_165 | spl10_217)), inference(forward_demodulation, [], [f1673, f1407])).
fof(f5797, plain, (~ spl10_15 | ~ spl10_40), inference(avatar_split_clause, [], [f5796, f808, f703])).
fof(f5796, plain, (~ (e4 = op(e4, e2)) | ~ spl10_40), inference(forward_demodulation, [], [f383, f810])).
fof(f5752, plain, (~ spl10_32 | ~ spl10_44 | spl10_165), inference(avatar_contradiction_clause, [], [f5751])).
fof(f5751, plain, ($false | (~ spl10_32 | ~ spl10_44 | spl10_165)), inference(subsumption_resolution, [], [f5744, f827])).
fof(f5744, plain, (~ (e3 = op(e3, e1)) | (~ spl10_32 | spl10_165)), inference(backward_demodulation, [], [f1406, f777])).
fof(f5702, plain, (~ spl10_51 | ~ spl10_75 | spl10_183), inference(avatar_split_clause, [], [f5701, f1502, f955, f855])).
fof(f5701, plain, (~ (e0 = op(e2, e4)) | (~ spl10_75 | spl10_183)), inference(forward_demodulation, [], [f1504, f957])).
fof(f5684, plain, (~ spl10_60 | ~ spl10_54 | spl10_162), inference(avatar_split_clause, [], [f5616, f1386, f867, f892])).
fof(f5616, plain, (~ (e4 = op(e2, e3)) | (~ spl10_54 | spl10_162)), inference(backward_demodulation, [], [f1388, f869])).
fof(f5683, plain, (spl10_40 | ~ spl10_8 | ~ spl10_166), inference(avatar_split_clause, [], [f5658, f1409, f674, f808])).
fof(f5658, plain, ((e4 = op(e3, e2)) | (~ spl10_8 | ~ spl10_166)), inference(backward_demodulation, [], [f1411, f676])).
fof(f5663, plain, (~ spl10_8 | ~ spl10_39 | ~ spl10_166), inference(avatar_contradiction_clause, [], [f5662])).
fof(f5662, plain, ($false | (~ spl10_8 | ~ spl10_39 | ~ spl10_166)), inference(subsumption_resolution, [], [f5661, f463])).
fof(f5661, plain, ((e3 = e4) | (~ spl10_8 | ~ spl10_39 | ~ spl10_166)), inference(forward_demodulation, [], [f5658, f806])).
fof(f5650, plain, (~ spl10_31 | ~ spl10_122 | spl10_216), inference(avatar_contradiction_clause, [], [f5649])).
fof(f5649, plain, ($false | (~ spl10_31 | ~ spl10_122 | spl10_216)), inference(subsumption_resolution, [], [f5643, f1155])).
fof(f5643, plain, (~ (op(e0, e0) = e1) | (~ spl10_31 | spl10_216)), inference(backward_demodulation, [], [f1669, f773])).
fof(f5648, plain, (~ spl10_31 | ~ spl10_47 | spl10_203), inference(avatar_contradiction_clause, [], [f5647])).
fof(f5647, plain, ($false | (~ spl10_31 | ~ spl10_47 | spl10_203)), inference(subsumption_resolution, [], [f5642, f840])).
fof(f5642, plain, (~ (e1 = op(e3, e0)) | (~ spl10_31 | spl10_203)), inference(backward_demodulation, [], [f1599, f773])).
fof(f5627, plain, (~ spl10_47 | ~ spl10_48), inference(avatar_contradiction_clause, [], [f5626])).
fof(f5626, plain, ($false | (~ spl10_47 | ~ spl10_48)), inference(subsumption_resolution, [], [f5625, f458])).
fof(f5625, plain, ((e1 = e2) | (~ spl10_47 | ~ spl10_48)), inference(backward_demodulation, [], [f844, f840])).
fof(f5624, plain, (~ spl10_51 | ~ spl10_54), inference(avatar_contradiction_clause, [], [f5623])).
fof(f5623, plain, ($false | (~ spl10_51 | ~ spl10_54)), inference(subsumption_resolution, [], [f5622, f456])).
fof(f5622, plain, ((e0 = e3) | (~ spl10_51 | ~ spl10_54)), inference(backward_demodulation, [], [f869, f857])).
fof(f5597, plain, (~ spl10_54 | ~ spl10_65 | spl10_188), inference(avatar_split_clause, [], [f5593, f1529, f913, f867])).
fof(f5593, plain, (~ (e3 = op(e2, e4)) | (~ spl10_65 | spl10_188)), inference(backward_demodulation, [], [f1531, f915])).
fof(f5575, plain, (~ spl10_112 | ~ spl10_114), inference(avatar_contradiction_clause, [], [f5574])).
fof(f5574, plain, ($false | (~ spl10_112 | ~ spl10_114)), inference(subsumption_resolution, [], [f5573, f459])).
fof(f5573, plain, ((e1 = e3) | (~ spl10_112 | ~ spl10_114)), inference(backward_demodulation, [], [f1121, f1113])).
fof(f5553, plain, (spl10_51 | ~ spl10_75 | ~ spl10_183), inference(avatar_split_clause, [], [f5552, f1502, f955, f855])).
fof(f5552, plain, ((e0 = op(e2, e4)) | (~ spl10_75 | ~ spl10_183)), inference(forward_demodulation, [], [f1503, f957])).
fof(f5531, plain, (~ spl10_53 | ~ spl10_68), inference(avatar_split_clause, [], [f5530, f926, f863])).
fof(f5530, plain, (~ (e2 = op(e2, e4)) | ~ spl10_68), inference(forward_demodulation, [], [f430, f928])).
fof(f5514, plain, (~ spl10_7 | ~ spl10_44 | ~ spl10_166), inference(avatar_contradiction_clause, [], [f5513])).
fof(f5513, plain, ($false | (~ spl10_7 | ~ spl10_44 | ~ spl10_166)), inference(subsumption_resolution, [], [f5512, f463])).
fof(f5512, plain, ((e3 = e4) | (~ spl10_7 | ~ spl10_44 | ~ spl10_166)), inference(forward_demodulation, [], [f5510, f827])).
fof(f5510, plain, ((e4 = op(e3, e1)) | (~ spl10_7 | ~ spl10_166)), inference(backward_demodulation, [], [f1411, f672])).
fof(f5481, plain, (~ spl10_58 | ~ spl10_68), inference(avatar_split_clause, [], [f5476, f926, f884])).
fof(f5476, plain, (~ (e2 = op(e2, e3)) | ~ spl10_68), inference(backward_demodulation, [], [f429, f928])).
fof(f5474, plain, (~ spl10_108 | spl10_114 | ~ spl10_142), inference(avatar_contradiction_clause, [], [f5473])).
fof(f5473, plain, ($false | (~ spl10_108 | spl10_114 | ~ spl10_142)), inference(subsumption_resolution, [], [f5469, f1120])).
fof(f1120, plain, (~ (e3 = op(e0, e2)) | spl10_114), inference(avatar_component_clause, [], [f1119])).
fof(f5469, plain, ((e3 = op(e0, e2)) | (~ spl10_108 | ~ spl10_142)), inference(backward_demodulation, [], [f1293, f1096])).
fof(f1293, plain, ((e3 = op(e0, op(e0, e3))) | ~ spl10_142), inference(avatar_component_clause, [], [f1292])).
fof(f5459, plain, (~ spl10_51 | ~ spl10_115 | spl10_147), inference(avatar_split_clause, [], [f5458, f1315, f1123, f855])).
fof(f5458, plain, (~ (e0 = op(e2, e4)) | (~ spl10_115 | spl10_147)), inference(forward_demodulation, [], [f1316, f1125])).
fof(f5411, plain, (~ spl10_116 | ~ spl10_122 | spl10_153), inference(avatar_contradiction_clause, [], [f5410])).
fof(f5410, plain, ($false | (~ spl10_116 | ~ spl10_122 | spl10_153)), inference(subsumption_resolution, [], [f5409, f1155])).
fof(f5409, plain, (~ (op(e0, e0) = e1) | (~ spl10_116 | spl10_153)), inference(forward_demodulation, [], [f1346, f1130])).
fof(f5406, plain, (~ spl10_96 | ~ spl10_116 | spl10_152), inference(avatar_contradiction_clause, [], [f5405])).
fof(f5405, plain, ($false | (~ spl10_96 | ~ spl10_116 | spl10_152)), inference(subsumption_resolution, [], [f5404, f1046])).
fof(f5404, plain, (~ (e0 = op(e1, e0)) | (~ spl10_116 | spl10_152)), inference(forward_demodulation, [], [f1340, f1130])).
fof(f5376, plain, (spl10_58 | ~ spl10_64 | ~ spl10_169), inference(avatar_split_clause, [], [f5270, f1429, f909, f884])).
fof(f5270, plain, ((e2 = op(e2, e3)) | (~ spl10_64 | ~ spl10_169)), inference(backward_demodulation, [], [f1431, f911])).
fof(f5333, plain, (~ spl10_31 | ~ spl10_35), inference(avatar_contradiction_clause, [], [f5332])).
fof(f5332, plain, ($false | (~ spl10_31 | ~ spl10_35)), inference(subsumption_resolution, [], [f5331, f457])).
fof(f5331, plain, ((e0 = e4) | (~ spl10_31 | ~ spl10_35)), inference(backward_demodulation, [], [f789, f773])).
fof(f5327, plain, (~ spl10_3 | ~ spl10_35 | spl10_232), inference(avatar_split_clause, [], [f5323, f1755, f787, f653])).
fof(f5323, plain, (~ (e2 = op(e4, e4)) | (~ spl10_35 | spl10_232)), inference(backward_demodulation, [], [f1757, f789])).
fof(f5326, plain, (spl10_29 | ~ spl10_35 | ~ spl10_165), inference(avatar_contradiction_clause, [], [f5325])).
fof(f5325, plain, ($false | (spl10_29 | ~ spl10_35 | ~ spl10_165)), inference(subsumption_resolution, [], [f5321, f763])).
fof(f763, plain, (~ (e3 = op(e3, e4)) | spl10_29), inference(avatar_component_clause, [], [f762])).
fof(f5321, plain, ((e3 = op(e3, e4)) | (~ spl10_35 | ~ spl10_165)), inference(backward_demodulation, [], [f1407, f789])).
fof(f5295, plain, (~ spl10_23 | ~ spl10_48), inference(avatar_split_clause, [], [f5286, f842, f737])).
fof(f5286, plain, (~ (e2 = op(e4, e0)) | ~ spl10_48), inference(backward_demodulation, [], [f363, f844])).
fof(f5285, plain, (~ spl10_51 | ~ spl10_52), inference(avatar_contradiction_clause, [], [f5284])).
fof(f5284, plain, ($false | (~ spl10_51 | ~ spl10_52)), inference(subsumption_resolution, [], [f5283, f454])).
fof(f5283, plain, ((e0 = e1) | (~ spl10_51 | ~ spl10_52)), inference(backward_demodulation, [], [f861, f857])).
fof(f5273, plain, (~ spl10_54 | ~ spl10_64), inference(avatar_split_clause, [], [f5268, f909, f867])).
fof(f5268, plain, (~ (e3 = op(e2, e4)) | ~ spl10_64), inference(backward_demodulation, [], [f432, f911])).
fof(f5264, plain, (~ spl10_75 | spl10_103 | ~ spl10_180), inference(avatar_contradiction_clause, [], [f5263])).
fof(f5263, plain, ($false | (~ spl10_75 | spl10_103 | ~ spl10_180)), inference(subsumption_resolution, [], [f5261, f1074])).
fof(f5261, plain, ((e2 = op(e0, e4)) | (~ spl10_75 | ~ spl10_180)), inference(backward_demodulation, [], [f1491, f957])).
fof(f5250, plain, (~ spl10_84 | spl10_139), inference(avatar_contradiction_clause, [], [f5249])).
fof(f5249, plain, ($false | (~ spl10_84 | spl10_139)), inference(subsumption_resolution, [], [f5245, f995])).
fof(f5245, plain, (~ (e3 = op(e1, e3)) | (~ spl10_84 | spl10_139)), inference(backward_demodulation, [], [f1280, f995])).
fof(f5239, plain, (~ spl10_38 | ~ spl10_88), inference(avatar_split_clause, [], [f5231, f1010, f800])).
fof(f5231, plain, (~ (e2 = op(e3, e2)) | ~ spl10_88), inference(backward_demodulation, [], [f379, f1012])).
fof(f5220, plain, (~ spl10_91 | ~ spl10_92 | spl10_198), inference(avatar_split_clause, [], [f5211, f1575, f1027, f1023])).
fof(f5211, plain, (~ (e0 = op(e1, e1)) | (~ spl10_92 | spl10_198)), inference(backward_demodulation, [], [f1577, f1029])).
fof(f5219, plain, (~ spl10_82 | ~ spl10_92), inference(avatar_split_clause, [], [f5206, f1027, f985])).
fof(f5206, plain, (~ (e1 = op(e1, e3)) | ~ spl10_92), inference(backward_demodulation, [], [f419, f1029])).
fof(f5218, plain, (~ spl10_17 | ~ spl10_92), inference(avatar_split_clause, [], [f5204, f1027, f712])).
fof(f5204, plain, (~ (e1 = op(e4, e1)) | ~ spl10_92), inference(backward_demodulation, [], [f370, f1029])).
fof(f5201, plain, (~ spl10_46 | ~ spl10_96), inference(avatar_split_clause, [], [f5189, f1044, f834])).
fof(f5189, plain, (~ (e0 = op(e3, e0)) | ~ spl10_96), inference(backward_demodulation, [], [f359, f1046])).
fof(f5176, plain, (~ spl10_107 | ~ spl10_116 | ~ spl10_142), inference(avatar_contradiction_clause, [], [f5175])).
fof(f5175, plain, ($false | (~ spl10_107 | ~ spl10_116 | ~ spl10_142)), inference(subsumption_resolution, [], [f5174, f456])).
fof(f5174, plain, ((e0 = e3) | (~ spl10_107 | ~ spl10_116 | ~ spl10_142)), inference(forward_demodulation, [], [f5169, f1130])).
fof(f5169, plain, ((e3 = op(e0, e1)) | (~ spl10_107 | ~ spl10_142)), inference(backward_demodulation, [], [f1293, f1092])).
fof(f5161, plain, (~ spl10_116 | ~ spl10_118), inference(avatar_contradiction_clause, [], [f5160])).
fof(f5160, plain, ($false | (~ spl10_116 | ~ spl10_118)), inference(subsumption_resolution, [], [f5159, f455])).
fof(f5159, plain, ((e0 = e2) | (~ spl10_116 | ~ spl10_118)), inference(backward_demodulation, [], [f1138, f1130])).
fof(f5154, plain, (~ spl10_70 | ~ spl10_118 | ~ spl10_122 | spl10_229), inference(avatar_contradiction_clause, [], [f5153])).
fof(f5153, plain, ($false | (~ spl10_70 | ~ spl10_118 | ~ spl10_122 | spl10_229)), inference(subsumption_resolution, [], [f5152, f936])).
fof(f5152, plain, (~ (e4 = op(e2, e1)) | (~ spl10_118 | ~ spl10_122 | spl10_229)), inference(forward_demodulation, [], [f5142, f1138])).
fof(f5142, plain, (~ (e4 = op(op(e0, e1), e1)) | (~ spl10_122 | spl10_229)), inference(backward_demodulation, [], [f1737, f1155])).
fof(f5150, plain, (~ spl10_92 | ~ spl10_122 | spl10_220), inference(avatar_split_clause, [], [f5140, f1687, f1153, f1027])).
fof(f5140, plain, (~ (e1 = op(e1, e1)) | (~ spl10_122 | spl10_220)), inference(backward_demodulation, [], [f1689, f1155])).
fof(f5092, plain, (spl10_224 | ~ spl10_165 | ~ spl10_217), inference(avatar_split_clause, [], [f5091, f1671, f1405, f1705])).
fof(f5091, plain, ((e0 = op(e3, op(e3, e3))) | (~ spl10_165 | ~ spl10_217)), inference(forward_demodulation, [], [f1672, f1407])).
fof(f5049, plain, (~ spl10_91 | ~ spl10_100 | spl10_201), inference(avatar_contradiction_clause, [], [f5048])).
fof(f5048, plain, ($false | (~ spl10_91 | ~ spl10_100 | spl10_201)), inference(subsumption_resolution, [], [f5047, f1062])).
fof(f5047, plain, (~ (e4 = op(e1, e0)) | (~ spl10_91 | spl10_201)), inference(forward_demodulation, [], [f1590, f1025])).
fof(f5032, plain, (~ spl10_1 | ~ spl10_23 | spl10_196), inference(avatar_contradiction_clause, [], [f5031])).
fof(f5031, plain, ($false | (~ spl10_1 | ~ spl10_23 | spl10_196)), inference(subsumption_resolution, [], [f5025, f739])).
fof(f5025, plain, (~ (e2 = op(e4, e0)) | (~ spl10_1 | spl10_196)), inference(backward_demodulation, [], [f1567, f647])).
fof(f4988, plain, (~ spl10_23 | ~ spl10_91 | ~ spl10_100 | spl10_199), inference(avatar_contradiction_clause, [], [f4987])).
fof(f4987, plain, ($false | (~ spl10_23 | ~ spl10_91 | ~ spl10_100 | spl10_199)), inference(subsumption_resolution, [], [f4986, f739])).
fof(f4986, plain, (~ (e2 = op(e4, e0)) | (~ spl10_91 | ~ spl10_100 | spl10_199)), inference(forward_demodulation, [], [f4978, f1062])).
fof(f4978, plain, (~ (e2 = op(op(e1, e0), e0)) | (~ spl10_91 | spl10_199)), inference(backward_demodulation, [], [f1581, f1025])).
fof(f4985, plain, (~ spl10_86 | ~ spl10_91), inference(avatar_split_clause, [], [f4972, f1023, f1002])).
fof(f4972, plain, (~ (e0 = op(e1, e2)) | ~ spl10_91), inference(backward_demodulation, [], [f418, f1025])).
fof(f4965, plain, (~ spl10_86 | ~ spl10_118 | spl10_152), inference(avatar_split_clause, [], [f4964, f1339, f1136, f1002])).
fof(f4964, plain, (~ (e0 = op(e1, e2)) | (~ spl10_118 | spl10_152)), inference(forward_demodulation, [], [f1340, f1138])).
fof(f4948, plain, (~ spl10_90 | ~ spl10_115), inference(avatar_split_clause, [], [f4947, f1123, f1018])).
fof(f4947, plain, (~ (e4 = op(e1, e2)) | ~ spl10_115), inference(forward_demodulation, [], [f374, f1125])).
fof(f4936, plain, (~ spl10_228 | ~ spl10_169 | spl10_215), inference(avatar_split_clause, [], [f4935, f1661, f1429, f1727])).
fof(f4935, plain, (~ (e0 = op(e2, op(e2, e2))) | (~ spl10_169 | spl10_215)), inference(forward_demodulation, [], [f1663, f1431])).
fof(f4928, plain, (spl10_64 | ~ spl10_38 | ~ spl10_170), inference(avatar_split_clause, [], [f4901, f1433, f800, f909])).
fof(f4901, plain, ((e3 = op(e2, e2)) | (~ spl10_38 | ~ spl10_170)), inference(forward_demodulation, [], [f1435, f802])).
fof(f802, plain, ((e2 = op(e3, e2)) | ~ spl10_38), inference(avatar_component_clause, [], [f800])).
fof(f4915, plain, (~ spl10_106 | ~ spl10_124 | spl10_142), inference(avatar_contradiction_clause, [], [f4914])).
fof(f4914, plain, ($false | (~ spl10_106 | ~ spl10_124 | spl10_142)), inference(subsumption_resolution, [], [f4913, f1163])).
fof(f4913, plain, (~ (op(e0, e0) = e3) | (~ spl10_106 | spl10_142)), inference(forward_demodulation, [], [f1294, f1088])).
fof(f4878, plain, (~ spl10_16 | spl10_22 | ~ spl10_179), inference(avatar_contradiction_clause, [], [f4877])).
fof(f4877, plain, ($false | (~ spl10_16 | spl10_22 | ~ spl10_179)), inference(subsumption_resolution, [], [f4876, f734])).
fof(f734, plain, (~ (e1 = op(e4, e0)) | spl10_22), inference(avatar_component_clause, [], [f733])).
fof(f4876, plain, ((e1 = op(e4, e0)) | (~ spl10_16 | ~ spl10_179)), inference(backward_demodulation, [], [f1482, f710])).
fof(f4845, plain, (~ spl10_38 | spl10_172), inference(avatar_contradiction_clause, [], [f4844])).
fof(f4844, plain, ($false | (~ spl10_38 | spl10_172)), inference(subsumption_resolution, [], [f4840, f802])).
fof(f4840, plain, (~ (e2 = op(e3, e2)) | (~ spl10_38 | spl10_172)), inference(backward_demodulation, [], [f1445, f802])).
fof(f4841, plain, (~ spl10_13 | ~ spl10_38), inference(avatar_split_clause, [], [f4834, f800, f695])).
fof(f4834, plain, (~ (e2 = op(e4, e2)) | ~ spl10_38), inference(backward_demodulation, [], [f383, f802])).
fof(f4833, plain, (~ spl10_17 | ~ spl10_42), inference(avatar_split_clause, [], [f4826, f817, f712])).
fof(f4826, plain, (~ (e1 = op(e4, e1)) | ~ spl10_42), inference(backward_demodulation, [], [f373, f819])).
fof(f4825, plain, (~ spl10_46 | spl10_184), inference(avatar_contradiction_clause, [], [f4824])).
fof(f4824, plain, ($false | (~ spl10_46 | spl10_184)), inference(subsumption_resolution, [], [f4821, f836])).
fof(f4821, plain, (~ (e0 = op(e3, e0)) | (~ spl10_46 | spl10_184)), inference(backward_demodulation, [], [f1509, f836])).
fof(f4822, plain, (~ spl10_21 | ~ spl10_46), inference(avatar_split_clause, [], [f4811, f834, f729])).
fof(f4811, plain, (~ (e0 = op(e4, e0)) | ~ spl10_46), inference(backward_demodulation, [], [f363, f836])).
fof(f4783, plain, (~ spl10_66 | ~ spl10_72 | spl10_183), inference(avatar_split_clause, [], [f4778, f1502, f943, f918])).
fof(f4778, plain, (~ (e0 = op(e2, e1)) | (~ spl10_72 | spl10_183)), inference(backward_demodulation, [], [f1504, f945])).
fof(f4779, plain, (~ spl10_47 | ~ spl10_72), inference(avatar_split_clause, [], [f4770, f943, f838])).
fof(f4770, plain, (~ (e1 = op(e3, e0)) | ~ spl10_72), inference(backward_demodulation, [], [f361, f945])).
fof(f4762, plain, (~ spl10_42 | ~ spl10_82 | spl10_138), inference(avatar_split_clause, [], [f4757, f1273, f985, f817])).
fof(f4757, plain, (~ (e1 = op(e3, e1)) | (~ spl10_82 | spl10_138)), inference(backward_demodulation, [], [f1274, f987])).
fof(f4761, plain, (~ spl10_77 | ~ spl10_82), inference(avatar_split_clause, [], [f4753, f985, f964])).
fof(f4753, plain, (~ (e1 = op(e1, e4)) | ~ spl10_82), inference(backward_demodulation, [], [f423, f987])).
fof(f4749, plain, (spl10_76 | ~ spl10_100 | ~ spl10_156), inference(avatar_contradiction_clause, [], [f4748])).
fof(f4748, plain, ($false | (spl10_76 | ~ spl10_100 | ~ spl10_156)), inference(subsumption_resolution, [], [f4746, f961])).
fof(f961, plain, (~ (e0 = op(e1, e4)) | spl10_76), inference(avatar_component_clause, [], [f960])).
fof(f4746, plain, ((e0 = op(e1, e4)) | (~ spl10_100 | ~ spl10_156)), inference(backward_demodulation, [], [f1359, f1062])).
fof(f1359, plain, ((e0 = op(e1, op(e1, e0))) | ~ spl10_156), inference(avatar_component_clause, [], [f1358])).
fof(f4738, plain, (~ spl10_77 | ~ spl10_102), inference(avatar_split_clause, [], [f4731, f1069, f964])).
fof(f4731, plain, (~ (e1 = op(e1, e4)) | ~ spl10_102), inference(backward_demodulation, [], [f394, f1071])).
fof(f4728, plain, (~ spl10_101 | ~ spl10_106), inference(avatar_split_clause, [], [f4720, f1086, f1065])).
fof(f4720, plain, (~ (e0 = op(e0, e4)) | ~ spl10_106), inference(backward_demodulation, [], [f413, f1088])).
fof(f4727, plain, (~ spl10_31 | ~ spl10_106), inference(avatar_split_clause, [], [f4718, f1086, f771])).
fof(f4718, plain, (~ (e0 = op(e3, e3)) | ~ spl10_106), inference(backward_demodulation, [], [f386, f1088])).
fof(f4707, plain, (~ spl10_112 | ~ spl10_118 | spl10_153), inference(avatar_split_clause, [], [f4702, f1344, f1136, f1111])).
fof(f4702, plain, (~ (e1 = op(e0, e2)) | (~ spl10_118 | spl10_153)), inference(backward_demodulation, [], [f1346, f1138])).
fof(f4651, plain, (~ spl10_98 | ~ spl10_86 | spl10_145), inference(avatar_split_clause, [], [f4650, f1306, f1002, f1052])).
fof(f4650, plain, (~ (e2 = op(e1, e0)) | (~ spl10_86 | spl10_145)), inference(forward_demodulation, [], [f1308, f1004])).
fof(f4649, plain, (spl10_34 | ~ spl10_94 | ~ spl10_238), inference(avatar_split_clause, [], [f4648, f1801, f1035, f783])).
fof(f4648, plain, ((e3 = op(e3, e3)) | (~ spl10_94 | ~ spl10_238)), inference(forward_demodulation, [], [f1802, f1037])).
fof(f1802, plain, ((e3 = op(op(e1, e1), op(e1, e1))) | ~ spl10_238), inference(avatar_component_clause, [], [f1801])).
fof(f4638, plain, (spl10_78 | ~ spl10_70 | ~ spl10_174), inference(avatar_split_clause, [], [f4553, f1457, f934, f968])).
fof(f4553, plain, ((e2 = op(e1, e4)) | (~ spl10_70 | ~ spl10_174)), inference(backward_demodulation, [], [f1459, f936])).
fof(f4633, plain, (~ spl10_58 | ~ spl10_64 | spl10_169), inference(avatar_split_clause, [], [f4279, f1429, f909, f884])).
fof(f4279, plain, (~ (e2 = op(e2, e3)) | (~ spl10_64 | spl10_169)), inference(backward_demodulation, [], [f1430, f911])).
fof(f4607, plain, (~ spl10_104 | ~ spl10_4), inference(avatar_split_clause, [], [f4606, f657, f1077])).
fof(f4606, plain, (~ (e3 = op(e0, e4)) | ~ spl10_4), inference(forward_demodulation, [], [f397, f659])).
fof(f4599, plain, (~ spl10_31 | ~ spl10_47 | ~ spl10_98 | spl10_202), inference(avatar_contradiction_clause, [], [f4598])).
fof(f4598, plain, ($false | (~ spl10_31 | ~ spl10_47 | ~ spl10_98 | spl10_202)), inference(subsumption_resolution, [], [f4597, f1054])).
fof(f4597, plain, (~ (e2 = op(e1, e0)) | (~ spl10_31 | ~ spl10_47 | spl10_202)), inference(forward_demodulation, [], [f4593, f840])).
fof(f4593, plain, (~ (e2 = op(op(e3, e0), e0)) | (~ spl10_31 | spl10_202)), inference(backward_demodulation, [], [f1595, f773])).
fof(f4579, plain, (~ spl10_57 | ~ spl10_64 | ~ spl10_85 | spl10_212), inference(avatar_contradiction_clause, [], [f4578])).
fof(f4578, plain, ($false | (~ spl10_57 | ~ spl10_64 | ~ spl10_85 | spl10_212)), inference(subsumption_resolution, [], [f4574, f999])).
fof(f4574, plain, (~ (e4 = op(e1, e3)) | (~ spl10_57 | ~ spl10_64 | spl10_212)), inference(backward_demodulation, [], [f4282, f882])).
fof(f4282, plain, (~ (e4 = op(op(e2, e3), e3)) | (~ spl10_64 | spl10_212)), inference(backward_demodulation, [], [f1647, f911])).
fof(f4565, plain, (~ spl10_66 | ~ spl10_70), inference(avatar_contradiction_clause, [], [f4564])).
fof(f4564, plain, ($false | (~ spl10_66 | ~ spl10_70)), inference(subsumption_resolution, [], [f4561, f457])).
fof(f4561, plain, ((e0 = e4) | (~ spl10_66 | ~ spl10_70)), inference(backward_demodulation, [], [f936, f920])).
fof(f4558, plain, (~ spl10_70 | ~ spl10_77 | ~ spl10_174), inference(avatar_contradiction_clause, [], [f4557])).
fof(f4557, plain, ($false | (~ spl10_70 | ~ spl10_77 | ~ spl10_174)), inference(subsumption_resolution, [], [f4556, f458])).
fof(f4556, plain, ((e1 = e2) | (~ spl10_70 | ~ spl10_77 | ~ spl10_174)), inference(forward_demodulation, [], [f4553, f966])).
fof(f966, plain, ((e1 = op(e1, e4)) | ~ spl10_77), inference(avatar_component_clause, [], [f964])).
fof(f4497, plain, (~ spl10_83 | ~ spl10_98), inference(avatar_split_clause, [], [f4496, f1052, f989])).
fof(f4496, plain, (~ (e2 = op(e1, e3)) | ~ spl10_98), inference(forward_demodulation, [], [f416, f1054])).
fof(f4478, plain, (~ spl10_41 | ~ spl10_43), inference(avatar_contradiction_clause, [], [f4477])).
fof(f4477, plain, ($false | (~ spl10_41 | ~ spl10_43)), inference(subsumption_resolution, [], [f4476, f455])).
fof(f4476, plain, ((e0 = e2) | (~ spl10_41 | ~ spl10_43)), inference(backward_demodulation, [], [f823, f815])).
fof(f4467, plain, (~ spl10_56 | ~ spl10_60), inference(avatar_contradiction_clause, [], [f4466])).
fof(f4466, plain, ($false | (~ spl10_56 | ~ spl10_60)), inference(subsumption_resolution, [], [f4465, f457])).
fof(f4465, plain, ((e0 = e4) | (~ spl10_56 | ~ spl10_60)), inference(backward_demodulation, [], [f894, f878])).
fof(f4461, plain, (~ spl10_67 | spl10_177), inference(avatar_contradiction_clause, [], [f4460])).
fof(f4460, plain, ($false | (~ spl10_67 | spl10_177)), inference(subsumption_resolution, [], [f4457, f924])).
fof(f4457, plain, (~ (e1 = op(e2, e1)) | (~ spl10_67 | spl10_177)), inference(backward_demodulation, [], [f1472, f924])).
fof(f4459, plain, (~ spl10_67 | spl10_93 | ~ spl10_174), inference(avatar_contradiction_clause, [], [f4458])).
fof(f4458, plain, ($false | (~ spl10_67 | spl10_93 | ~ spl10_174)), inference(subsumption_resolution, [], [f4456, f1032])).
fof(f1032, plain, (~ (e2 = op(e1, e1)) | spl10_93), inference(avatar_component_clause, [], [f1031])).
fof(f4456, plain, ((e2 = op(e1, e1)) | (~ spl10_67 | ~ spl10_174)), inference(backward_demodulation, [], [f1459, f924])).
fof(f4452, plain, (~ spl10_94 | ~ spl10_95), inference(avatar_contradiction_clause, [], [f4451])).
fof(f4451, plain, ($false | (~ spl10_94 | ~ spl10_95)), inference(subsumption_resolution, [], [f4450, f463])).
fof(f4450, plain, ((e3 = e4) | (~ spl10_94 | ~ spl10_95)), inference(backward_demodulation, [], [f1041, f1037])).
fof(f4449, plain, (~ spl10_86 | ~ spl10_98 | spl10_156), inference(avatar_contradiction_clause, [], [f4448])).
fof(f4448, plain, ($false | (~ spl10_86 | ~ spl10_98 | spl10_156)), inference(subsumption_resolution, [], [f4447, f1004])).
fof(f4447, plain, (~ (e0 = op(e1, e2)) | (~ spl10_98 | spl10_156)), inference(backward_demodulation, [], [f1360, f1054])).
fof(f4442, plain, (~ spl10_112 | ~ spl10_113), inference(avatar_contradiction_clause, [], [f4441])).
fof(f4441, plain, ($false | (~ spl10_112 | ~ spl10_113)), inference(subsumption_resolution, [], [f4440, f458])).
fof(f4440, plain, ((e1 = e2) | (~ spl10_112 | ~ spl10_113)), inference(forward_demodulation, [], [f1117, f1113])).
fof(f4406, plain, (spl10_125 | ~ spl10_21 | ~ spl10_182), inference(avatar_split_clause, [], [f4179, f1497, f729, f1165])).
fof(f4179, plain, ((op(e0, e0) = e4) | (~ spl10_21 | ~ spl10_182)), inference(forward_demodulation, [], [f1499, f731])).
fof(f4382, plain, (spl10_98 | ~ spl10_86 | ~ spl10_145), inference(avatar_split_clause, [], [f4242, f1306, f1002, f1052])).
fof(f4242, plain, ((e2 = op(e1, e0)) | (~ spl10_86 | ~ spl10_145)), inference(backward_demodulation, [], [f1307, f1004])).
fof(f4375, plain, (~ spl10_70 | ~ spl10_32 | ~ spl10_43 | spl10_213), inference(avatar_split_clause, [], [f4314, f1651, f821, f775, f934])).
fof(f4314, plain, (~ (e4 = op(e2, e1)) | (~ spl10_32 | ~ spl10_43 | spl10_213)), inference(forward_demodulation, [], [f4308, f823])).
fof(f4308, plain, (~ (e4 = op(op(e3, e1), e1)) | (~ spl10_32 | spl10_213)), inference(backward_demodulation, [], [f1653, f777])).
fof(f4374, plain, (~ spl10_66 | ~ spl10_32 | ~ spl10_43 | spl10_217), inference(avatar_split_clause, [], [f4316, f1671, f821, f775, f918])).
fof(f4316, plain, (~ (e0 = op(e2, e1)) | (~ spl10_32 | ~ spl10_43 | spl10_217)), inference(forward_demodulation, [], [f4309, f823])).
fof(f4309, plain, (~ (e0 = op(op(e3, e1), e1)) | (~ spl10_32 | spl10_217)), inference(backward_demodulation, [], [f1673, f777])).
fof(f4339, plain, (~ spl10_6 | ~ spl10_9), inference(avatar_contradiction_clause, [], [f4338])).
fof(f4338, plain, ($false | (~ spl10_6 | ~ spl10_9)), inference(subsumption_resolution, [], [f4337, f456])).
fof(f4337, plain, ((e0 = e3) | (~ spl10_6 | ~ spl10_9)), inference(backward_demodulation, [], [f680, f668])).
fof(f4334, plain, (~ spl10_13 | spl10_173), inference(avatar_contradiction_clause, [], [f4333])).
fof(f4333, plain, ($false | (~ spl10_13 | spl10_173)), inference(subsumption_resolution, [], [f4329, f697])).
fof(f4329, plain, (~ (e2 = op(e4, e2)) | (~ spl10_13 | spl10_173)), inference(backward_demodulation, [], [f1451, f697])).
fof(f4332, plain, (~ spl10_13 | ~ spl10_64 | ~ spl10_171), inference(avatar_contradiction_clause, [], [f4331])).
fof(f4331, plain, ($false | (~ spl10_13 | ~ spl10_64 | ~ spl10_171)), inference(subsumption_resolution, [], [f4330, f463])).
fof(f4330, plain, ((e3 = e4) | (~ spl10_13 | ~ spl10_64 | ~ spl10_171)), inference(forward_demodulation, [], [f4328, f911])).
fof(f4328, plain, ((e4 = op(e2, e2)) | (~ spl10_13 | ~ spl10_171)), inference(backward_demodulation, [], [f1439, f697])).
fof(f4323, plain, (~ spl10_26 | ~ spl10_29), inference(avatar_contradiction_clause, [], [f4322])).
fof(f4322, plain, ($false | (~ spl10_26 | ~ spl10_29)), inference(subsumption_resolution, [], [f4321, f456])).
fof(f4321, plain, ((e0 = e3) | (~ spl10_26 | ~ spl10_29)), inference(backward_demodulation, [], [f764, f752])).
fof(f4285, plain, (~ spl10_14 | ~ spl10_64), inference(avatar_split_clause, [], [f4277, f909, f699])).
fof(f4277, plain, (~ (e3 = op(e4, e2)) | ~ spl10_64), inference(backward_demodulation, [], [f382, f911])).
fof(f4275, plain, (~ spl10_66 | ~ spl10_72 | spl10_177), inference(avatar_contradiction_clause, [], [f4274])).
fof(f4274, plain, ($false | (~ spl10_66 | ~ spl10_72 | spl10_177)), inference(subsumption_resolution, [], [f4270, f945])).
fof(f4270, plain, (~ (e1 = op(e2, e0)) | (~ spl10_66 | spl10_177)), inference(backward_demodulation, [], [f1472, f920])).
fof(f4243, plain, (~ spl10_36 | ~ spl10_86), inference(avatar_split_clause, [], [f4234, f1002, f792])).
fof(f4234, plain, (~ (e0 = op(e3, e2)) | ~ spl10_86), inference(backward_demodulation, [], [f379, f1004])).
fof(f4226, plain, (~ spl10_112 | ~ spl10_115), inference(avatar_contradiction_clause, [], [f4225])).
fof(f4225, plain, ($false | (~ spl10_112 | ~ spl10_115)), inference(subsumption_resolution, [], [f4224, f460])).
fof(f4224, plain, ((e1 = e4) | (~ spl10_112 | ~ spl10_115)), inference(backward_demodulation, [], [f1125, f1113])).
fof(f4210, plain, (~ spl10_115 | ~ spl10_123 | spl10_223), inference(avatar_contradiction_clause, [], [f4209])).
fof(f4209, plain, ($false | (~ spl10_115 | ~ spl10_123 | spl10_223)), inference(subsumption_resolution, [], [f4202, f1125])).
fof(f4202, plain, (~ (e4 = op(e0, e2)) | (~ spl10_123 | spl10_223)), inference(backward_demodulation, [], [f1702, f1159])).
fof(f4190, plain, (spl10_33 | ~ spl10_59 | ~ spl10_164), inference(avatar_contradiction_clause, [], [f4189])).
fof(f4189, plain, ($false | (spl10_33 | ~ spl10_59 | ~ spl10_164)), inference(subsumption_resolution, [], [f4188, f780])).
fof(f780, plain, (~ (e2 = op(e3, e3)) | spl10_33), inference(avatar_component_clause, [], [f779])).
fof(f4188, plain, ((e2 = op(e3, e3)) | (~ spl10_59 | ~ spl10_164)), inference(forward_demodulation, [], [f1403, f890])).
fof(f4187, plain, (~ spl10_59 | spl10_167), inference(avatar_contradiction_clause, [], [f4186])).
fof(f4186, plain, ($false | (~ spl10_59 | spl10_167)), inference(subsumption_resolution, [], [f4185, f890])).
fof(f4185, plain, (~ (e3 = op(e2, e3)) | (~ spl10_59 | spl10_167)), inference(forward_demodulation, [], [f1416, f890])).
fof(f4184, plain, (~ spl10_9 | spl10_168), inference(avatar_contradiction_clause, [], [f4183])).
fof(f4183, plain, ($false | (~ spl10_9 | spl10_168)), inference(subsumption_resolution, [], [f4182, f680])).
fof(f4182, plain, (~ (e3 = op(e4, e3)) | (~ spl10_9 | spl10_168)), inference(forward_demodulation, [], [f1423, f680])).
fof(f4178, plain, (spl10_26 | ~ spl10_50 | ~ spl10_184), inference(avatar_split_clause, [], [f4177, f1507, f850, f750])).
fof(f4177, plain, ((e0 = op(e3, e4)) | (~ spl10_50 | ~ spl10_184)), inference(forward_demodulation, [], [f1508, f852])).
fof(f1508, plain, ((e0 = op(e3, op(e3, e0))) | ~ spl10_184), inference(avatar_component_clause, [], [f1507])).
fof(f4149, plain, (~ spl10_21 | ~ spl10_101 | spl10_135), inference(avatar_contradiction_clause, [], [f4148])).
fof(f4148, plain, ($false | (~ spl10_21 | ~ spl10_101 | spl10_135)), inference(subsumption_resolution, [], [f4147, f731])).
fof(f4147, plain, (~ (e0 = op(e4, e0)) | (~ spl10_101 | spl10_135)), inference(forward_demodulation, [], [f1260, f1067])).
fof(f4106, plain, (~ spl10_21 | spl10_185), inference(avatar_contradiction_clause, [], [f4105])).
fof(f4105, plain, ($false | (~ spl10_21 | spl10_185)), inference(subsumption_resolution, [], [f4104, f731])).
fof(f4104, plain, (~ (e0 = op(e4, e0)) | (~ spl10_21 | spl10_185)), inference(forward_demodulation, [], [f1515, f731])).
fof(f4094, plain, (~ spl10_17 | spl10_179), inference(avatar_contradiction_clause, [], [f4093])).
fof(f4093, plain, ($false | (~ spl10_17 | spl10_179)), inference(subsumption_resolution, [], [f4092, f714])).
fof(f714, plain, ((e1 = op(e4, e1)) | ~ spl10_17), inference(avatar_component_clause, [], [f712])).
fof(f4092, plain, (~ (e1 = op(e4, e1)) | (~ spl10_17 | spl10_179)), inference(forward_demodulation, [], [f1483, f714])).
fof(f4063, plain, (~ spl10_8 | ~ spl10_9), inference(avatar_contradiction_clause, [], [f4062])).
fof(f4062, plain, ($false | (~ spl10_8 | ~ spl10_9)), inference(subsumption_resolution, [], [f4061, f461])).
fof(f4061, plain, ((e2 = e3) | (~ spl10_8 | ~ spl10_9)), inference(forward_demodulation, [], [f680, f676])).
fof(f4042, plain, (~ spl10_28 | ~ spl10_35 | ~ spl10_51 | spl10_217), inference(avatar_contradiction_clause, [], [f4041])).
fof(f4041, plain, ($false | (~ spl10_28 | ~ spl10_35 | ~ spl10_51 | spl10_217)), inference(subsumption_resolution, [], [f4040, f857])).
fof(f4040, plain, (~ (e0 = op(e2, e4)) | (~ spl10_28 | ~ spl10_35 | spl10_217)), inference(forward_demodulation, [], [f4033, f760])).
fof(f4033, plain, (~ (e0 = op(op(e3, e4), e4)) | (~ spl10_35 | spl10_217)), inference(backward_demodulation, [], [f1673, f789])).
fof(f4023, plain, (~ spl10_47 | ~ spl10_50), inference(avatar_contradiction_clause, [], [f4022])).
fof(f4022, plain, ($false | (~ spl10_47 | ~ spl10_50)), inference(subsumption_resolution, [], [f4021, f460])).
fof(f4021, plain, ((e1 = e4) | (~ spl10_47 | ~ spl10_50)), inference(backward_demodulation, [], [f852, f840])).
fof(f4018, plain, (~ spl10_51 | ~ spl10_53), inference(avatar_contradiction_clause, [], [f4017])).
fof(f4017, plain, ($false | (~ spl10_51 | ~ spl10_53)), inference(subsumption_resolution, [], [f4016, f455])).
fof(f4016, plain, ((e0 = e2) | (~ spl10_51 | ~ spl10_53)), inference(backward_demodulation, [], [f865, f857])).
fof(f4002, plain, (~ spl10_41 | ~ spl10_62 | ~ spl10_69 | spl10_215), inference(avatar_split_clause, [], [f4001, f1661, f930, f901, f813])).
fof(f4001, plain, (~ (e0 = op(e3, e1)) | (~ spl10_62 | ~ spl10_69 | spl10_215)), inference(forward_demodulation, [], [f3995, f932])).
fof(f3995, plain, (~ (e0 = op(op(e2, e1), e1)) | (~ spl10_62 | spl10_215)), inference(backward_demodulation, [], [f1663, f903])).
fof(f3952, plain, (spl10_81 | ~ spl10_99 | ~ spl10_156), inference(avatar_split_clause, [], [f3950, f1358, f1056, f981])).
fof(f3950, plain, ((e0 = op(e1, e3)) | (~ spl10_99 | ~ spl10_156)), inference(backward_demodulation, [], [f1359, f1058])).
fof(f3943, plain, (~ spl10_101 | ~ spl10_102), inference(avatar_contradiction_clause, [], [f3942])).
fof(f3942, plain, ($false | (~ spl10_101 | ~ spl10_102)), inference(subsumption_resolution, [], [f3941, f454])).
fof(f3941, plain, ((e0 = e1) | (~ spl10_101 | ~ spl10_102)), inference(backward_demodulation, [], [f1071, f1067])).
fof(f3907, plain, (spl10_5 | ~ spl10_95 | ~ spl10_242), inference(avatar_split_clause, [], [f3906, f1841, f1039, f661])).
fof(f3906, plain, ((e4 = op(e4, e4)) | (~ spl10_95 | ~ spl10_242)), inference(forward_demodulation, [], [f1842, f1041])).
fof(f3853, plain, (~ spl10_28 | ~ spl10_130), inference(avatar_contradiction_clause, [], [f3852])).
fof(f3852, plain, ($false | (~ spl10_28 | ~ spl10_130)), inference(subsumption_resolution, [], [f3851, f461])).
fof(f3851, plain, ((e2 = e3) | (~ spl10_28 | ~ spl10_130)), inference(forward_demodulation, [], [f3668, f760])).
fof(f3668, plain, ((e3 = op(e3, e4)) | ~ spl10_130), inference(backward_demodulation, [], [f300, f1188])).
fof(f3834, plain, (~ spl10_13 | ~ spl10_88), inference(avatar_split_clause, [], [f3833, f1010, f695])).
fof(f3833, plain, (~ (e2 = op(e4, e2)) | ~ spl10_88), inference(forward_demodulation, [], [f380, f1012])).
fof(f3831, plain, (~ spl10_8 | ~ spl10_130), inference(avatar_contradiction_clause, [], [f3830])).
fof(f3830, plain, ($false | (~ spl10_8 | ~ spl10_130)), inference(subsumption_resolution, [], [f3829, f461])).
fof(f3829, plain, ((e2 = e3) | (~ spl10_8 | ~ spl10_130)), inference(forward_demodulation, [], [f3667, f676])).
fof(f3667, plain, ((e3 = op(e4, e3)) | ~ spl10_130), inference(backward_demodulation, [], [f299, f1188])).
fof(f3812, plain, (~ spl10_18 | ~ spl10_19), inference(avatar_contradiction_clause, [], [f3811])).
fof(f3811, plain, ($false | (~ spl10_18 | ~ spl10_19)), inference(subsumption_resolution, [], [f3810, f461])).
fof(f3810, plain, ((e2 = e3) | (~ spl10_18 | ~ spl10_19)), inference(backward_demodulation, [], [f722, f718])).
fof(f3796, plain, (~ spl10_36 | ~ spl10_40), inference(avatar_contradiction_clause, [], [f3795])).
fof(f3795, plain, ($false | (~ spl10_36 | ~ spl10_40)), inference(subsumption_resolution, [], [f3794, f457])).
fof(f3794, plain, ((e0 = e4) | (~ spl10_36 | ~ spl10_40)), inference(backward_demodulation, [], [f810, f794])).
fof(f3777, plain, (~ spl10_56 | ~ spl10_58), inference(avatar_contradiction_clause, [], [f3776])).
fof(f3776, plain, ($false | (~ spl10_56 | ~ spl10_58)), inference(subsumption_resolution, [], [f3775, f455])).
fof(f3775, plain, ((e0 = e2) | (~ spl10_56 | ~ spl10_58)), inference(backward_demodulation, [], [f886, f878])).
fof(f3742, plain, (~ spl10_88 | spl10_145), inference(avatar_contradiction_clause, [], [f3741])).
fof(f3741, plain, ($false | (~ spl10_88 | spl10_145)), inference(subsumption_resolution, [], [f3735, f1012])).
fof(f3735, plain, (~ (e2 = op(e1, e2)) | (~ spl10_88 | spl10_145)), inference(backward_demodulation, [], [f1308, f1012])).
fof(f3693, plain, (~ spl10_123 | ~ spl10_124), inference(avatar_contradiction_clause, [], [f3692])).
fof(f3692, plain, ($false | (~ spl10_123 | ~ spl10_124)), inference(subsumption_resolution, [], [f3691, f461])).
fof(f3691, plain, ((e2 = e3) | (~ spl10_123 | ~ spl10_124)), inference(backward_demodulation, [], [f1163, f1159])).
fof(f3681, plain, (~ spl10_12 | ~ spl10_130), inference(avatar_contradiction_clause, [], [f3680])).
fof(f3680, plain, ($false | (~ spl10_12 | ~ spl10_130)), inference(subsumption_resolution, [], [f3679, f458])).
fof(f3679, plain, ((e1 = e2) | (~ spl10_12 | ~ spl10_130)), inference(forward_demodulation, [], [f3665, f693])).
fof(f3665, plain, ((e2 = op(e4, e2)) | ~ spl10_130), inference(backward_demodulation, [], [f297, f1188])).
fof(f3614, plain, (~ spl10_4 | ~ spl10_95 | spl10_238), inference(avatar_split_clause, [], [f3558, f1801, f1039, f657])).
fof(f3558, plain, (~ (e3 = op(e4, e4)) | (~ spl10_95 | spl10_238)), inference(backward_demodulation, [], [f1803, f1041])).
fof(f3603, plain, (~ spl10_26 | ~ spl10_30), inference(avatar_contradiction_clause, [], [f3602])).
fof(f3602, plain, ($false | (~ spl10_26 | ~ spl10_30)), inference(subsumption_resolution, [], [f3601, f457])).
fof(f3601, plain, ((e0 = e4) | (~ spl10_26 | ~ spl10_30)), inference(backward_demodulation, [], [f768, f752])).
fof(f3591, plain, (~ spl10_57 | ~ spl10_58), inference(avatar_contradiction_clause, [], [f3590])).
fof(f3590, plain, ($false | (~ spl10_57 | ~ spl10_58)), inference(subsumption_resolution, [], [f3589, f458])).
fof(f3589, plain, ((e1 = e2) | (~ spl10_57 | ~ spl10_58)), inference(backward_demodulation, [], [f886, f882])).
fof(f3588, plain, (~ spl10_61 | ~ spl10_64), inference(avatar_contradiction_clause, [], [f3587])).
fof(f3587, plain, ($false | (~ spl10_61 | ~ spl10_64)), inference(subsumption_resolution, [], [f3586, f456])).
fof(f3586, plain, ((e0 = e3) | (~ spl10_61 | ~ spl10_64)), inference(backward_demodulation, [], [f911, f899])).
fof(f3565, plain, (~ spl10_81 | ~ spl10_82), inference(avatar_contradiction_clause, [], [f3564])).
fof(f3564, plain, ($false | (~ spl10_81 | ~ spl10_82)), inference(subsumption_resolution, [], [f3563, f454])).
fof(f3563, plain, ((e0 = e1) | (~ spl10_81 | ~ spl10_82)), inference(backward_demodulation, [], [f987, f983])).
fof(f3548, plain, (~ spl10_34 | ~ spl10_124 | spl10_241), inference(avatar_contradiction_clause, [], [f3547])).
fof(f3547, plain, ($false | (~ spl10_34 | ~ spl10_124 | spl10_241)), inference(subsumption_resolution, [], [f3544, f785])).
fof(f3544, plain, (~ (e3 = op(e3, e3)) | (~ spl10_124 | spl10_241)), inference(backward_demodulation, [], [f1821, f1163])).
fof(f3533, plain, (~ spl10_38 | ~ spl10_128), inference(avatar_contradiction_clause, [], [f3532])).
fof(f3532, plain, ($false | (~ spl10_38 | ~ spl10_128)), inference(subsumption_resolution, [], [f3531, f461])).
fof(f3531, plain, ((e2 = e3) | (~ spl10_38 | ~ spl10_128)), inference(forward_demodulation, [], [f3512, f802])).
fof(f3512, plain, ((e3 = op(e3, e2)) | ~ spl10_128), inference(backward_demodulation, [], [f300, f1180])).
fof(f3530, plain, (~ spl10_58 | ~ spl10_128), inference(avatar_contradiction_clause, [], [f3529])).
fof(f3529, plain, ($false | (~ spl10_58 | ~ spl10_128)), inference(subsumption_resolution, [], [f3528, f461])).
fof(f3528, plain, ((e2 = e3) | (~ spl10_58 | ~ spl10_128)), inference(forward_demodulation, [], [f3511, f886])).
fof(f3511, plain, ((e3 = op(e2, e3)) | ~ spl10_128), inference(backward_demodulation, [], [f299, f1180])).
fof(f3458, plain, (~ spl10_49 | ~ spl10_34), inference(avatar_split_clause, [], [f3136, f783, f846])).
fof(f3136, plain, (~ (e3 = op(e3, e0)) | ~ spl10_34), inference(forward_demodulation, [], [f436, f785])).
fof(f3431, plain, (~ spl10_3 | ~ spl10_13), inference(avatar_split_clause, [], [f3428, f695, f653])).
fof(f3428, plain, (~ (e2 = op(e4, e4)) | ~ spl10_13), inference(backward_demodulation, [], [f452, f697])).
fof(f3427, plain, (~ spl10_16 | ~ spl10_17), inference(avatar_contradiction_clause, [], [f3426])).
fof(f3426, plain, ($false | (~ spl10_16 | ~ spl10_17)), inference(subsumption_resolution, [], [f3425, f454])).
fof(f3425, plain, ((e0 = e1) | (~ spl10_16 | ~ spl10_17)), inference(backward_demodulation, [], [f714, f710])).
fof(f3393, plain, (~ spl10_69 | ~ spl10_70), inference(avatar_contradiction_clause, [], [f3392])).
fof(f3392, plain, ($false | (~ spl10_69 | ~ spl10_70)), inference(subsumption_resolution, [], [f3391, f463])).
fof(f3391, plain, ((e3 = e4) | (~ spl10_69 | ~ spl10_70)), inference(forward_demodulation, [], [f936, f932])).
fof(f3390, plain, (~ spl10_52 | ~ spl10_72), inference(avatar_split_clause, [], [f3385, f943, f859])).
fof(f3385, plain, (~ (e1 = op(e2, e4)) | ~ spl10_72), inference(backward_demodulation, [], [f427, f945])).
fof(f3373, plain, (~ spl10_86 | ~ spl10_89), inference(avatar_contradiction_clause, [], [f3372])).
fof(f3372, plain, ($false | (~ spl10_86 | ~ spl10_89)), inference(subsumption_resolution, [], [f3371, f456])).
fof(f3371, plain, ((e0 = e3) | (~ spl10_86 | ~ spl10_89)), inference(backward_demodulation, [], [f1016, f1004])).
fof(f3357, plain, (~ spl10_61 | ~ spl10_93 | spl10_198), inference(avatar_split_clause, [], [f3350, f1575, f1031, f897])).
fof(f3350, plain, (~ (e0 = op(e2, e2)) | (~ spl10_93 | spl10_198)), inference(backward_demodulation, [], [f1577, f1033])).
fof(f3344, plain, (~ spl10_2 | ~ spl10_102), inference(avatar_split_clause, [], [f3339, f1069, f649])).
fof(f3339, plain, (~ (e1 = op(e4, e4)) | ~ spl10_102), inference(backward_demodulation, [], [f397, f1071])).
fof(f3311, plain, (spl10_102 | ~ spl10_100 | ~ spl10_155), inference(avatar_split_clause, [], [f3193, f1353, f1060, f1069])).
fof(f3193, plain, ((e1 = op(e0, e4)) | (~ spl10_100 | ~ spl10_155)), inference(backward_demodulation, [], [f1355, f1062])).
fof(f3291, plain, (~ spl10_5 | ~ spl10_30), inference(avatar_split_clause, [], [f3290, f766, f661])).
fof(f3290, plain, (~ (e4 = op(e4, e4)) | ~ spl10_30), inference(forward_demodulation, [], [f403, f768])).
fof(f3289, plain, (~ spl10_1 | ~ spl10_2), inference(avatar_contradiction_clause, [], [f3288])).
fof(f3288, plain, ($false | (~ spl10_1 | ~ spl10_2)), inference(subsumption_resolution, [], [f3287, f454])).
fof(f3287, plain, ((e0 = e1) | (~ spl10_1 | ~ spl10_2)), inference(backward_demodulation, [], [f651, f647])).
fof(f3286, plain, (~ spl10_2 | ~ spl10_3), inference(avatar_contradiction_clause, [], [f3285])).
fof(f3285, plain, ($false | (~ spl10_2 | ~ spl10_3)), inference(subsumption_resolution, [], [f3284, f458])).
fof(f3284, plain, ((e1 = e2) | (~ spl10_2 | ~ spl10_3)), inference(backward_demodulation, [], [f655, f651])).
fof(f3256, plain, (~ spl10_22 | ~ spl10_61 | ~ spl10_75 | spl10_187), inference(avatar_split_clause, [], [f3255, f1525, f955, f897, f733])).
fof(f3255, plain, (~ (e1 = op(e4, e0)) | (~ spl10_61 | ~ spl10_75 | spl10_187)), inference(forward_demodulation, [], [f3249, f957])).
fof(f3249, plain, (~ (e1 = op(op(e2, e0), e0)) | (~ spl10_61 | spl10_187)), inference(backward_demodulation, [], [f1527, f899])).
fof(f3111, plain, (~ spl10_68 | ~ spl10_18), inference(avatar_split_clause, [], [f3110, f716, f926])).
fof(f3110, plain, (~ (e2 = op(e2, e1)) | ~ spl10_18), inference(forward_demodulation, [], [f372, f718])).
fof(f3103, plain, (~ spl10_14 | ~ spl10_3 | spl10_197), inference(avatar_split_clause, [], [f3102, f1570, f653, f699])).
fof(f3102, plain, (~ (e3 = op(e4, e2)) | (~ spl10_3 | spl10_197)), inference(forward_demodulation, [], [f1572, f655])).
fof(f3098, plain, (~ spl10_11 | ~ spl10_3 | spl10_225), inference(avatar_split_clause, [], [f3097, f1711, f653, f687])).
fof(f3097, plain, (~ (e0 = op(e4, e2)) | (~ spl10_3 | spl10_225)), inference(forward_demodulation, [], [f1713, f655])).
fof(f3096, plain, (~ spl10_3 | ~ spl10_4), inference(avatar_contradiction_clause, [], [f3095])).
fof(f3095, plain, ($false | (~ spl10_3 | ~ spl10_4)), inference(subsumption_resolution, [], [f3094, f461])).
fof(f3094, plain, ((e2 = e3) | (~ spl10_3 | ~ spl10_4)), inference(backward_demodulation, [], [f659, f655])).
fof(f3085, plain, (~ spl10_5 | ~ spl10_10), inference(avatar_split_clause, [], [f3084, f682, f661])).
fof(f3084, plain, (~ (e4 = op(e4, e4)) | ~ spl10_10), inference(backward_demodulation, [], [f453, f684])).
fof(f3077, plain, (~ spl10_16 | ~ spl10_18), inference(avatar_contradiction_clause, [], [f3076])).
fof(f3076, plain, ($false | (~ spl10_16 | ~ spl10_18)), inference(subsumption_resolution, [], [f3075, f455])).
fof(f3075, plain, ((e0 = e2) | (~ spl10_16 | ~ spl10_18)), inference(forward_demodulation, [], [f718, f710])).
fof(f3039, plain, (~ spl10_51 | ~ spl10_55), inference(avatar_contradiction_clause, [], [f3038])).
fof(f3038, plain, ($false | (~ spl10_51 | ~ spl10_55)), inference(subsumption_resolution, [], [f3037, f457])).
fof(f3037, plain, ((e0 = e4) | (~ spl10_51 | ~ spl10_55)), inference(backward_demodulation, [], [f873, f857])).
fof(f3000, plain, (~ spl10_82 | ~ spl10_83), inference(avatar_contradiction_clause, [], [f2999])).
fof(f2999, plain, ($false | (~ spl10_82 | ~ spl10_83)), inference(subsumption_resolution, [], [f2998, f458])).
fof(f2998, plain, ((e1 = e2) | (~ spl10_82 | ~ spl10_83)), inference(backward_demodulation, [], [f991, f987])).
fof(f2984, plain, (~ spl10_99 | ~ spl10_106 | ~ spl10_155), inference(avatar_contradiction_clause, [], [f2983])).
fof(f2983, plain, ($false | (~ spl10_99 | ~ spl10_106 | ~ spl10_155)), inference(subsumption_resolution, [], [f2982, f454])).
fof(f2982, plain, ((e0 = e1) | (~ spl10_99 | ~ spl10_106 | ~ spl10_155)), inference(forward_demodulation, [], [f2979, f1088])).
fof(f2979, plain, ((e1 = op(e0, e3)) | (~ spl10_99 | ~ spl10_155)), inference(backward_demodulation, [], [f1355, f1058])).
fof(f2962, plain, (~ spl10_102 | ~ spl10_112), inference(avatar_split_clause, [], [f2959, f1111, f1069])).
fof(f2959, plain, (~ (e1 = op(e0, e4)) | ~ spl10_112), inference(backward_demodulation, [], [f412, f1113])).
fof(f2951, plain, (~ spl10_98 | ~ spl10_123), inference(avatar_split_clause, [], [f2943, f1157, f1052])).
fof(f2943, plain, (~ (e2 = op(e1, e0)) | ~ spl10_123), inference(backward_demodulation, [], [f354, f1159])).
fof(f2888, plain, (~ spl10_2 | ~ spl10_22), inference(avatar_split_clause, [], [f2887, f733, f649])).
fof(f2887, plain, (~ (e1 = op(e4, e4)) | ~ spl10_22), inference(forward_demodulation, [], [f447, f735])).
fof(f2852, plain, (~ spl10_56 | ~ spl10_57), inference(avatar_contradiction_clause, [], [f2851])).
fof(f2851, plain, ($false | (~ spl10_56 | ~ spl10_57)), inference(subsumption_resolution, [], [f2850, f454])).
fof(f2850, plain, ((e0 = e1) | (~ spl10_56 | ~ spl10_57)), inference(backward_demodulation, [], [f882, f878])).
fof(f2833, plain, (spl10_112 | ~ spl10_98 | ~ spl10_155), inference(avatar_split_clause, [], [f2829, f1353, f1052, f1111])).
fof(f2829, plain, ((e1 = op(e0, e2)) | (~ spl10_98 | ~ spl10_155)), inference(backward_demodulation, [], [f1355, f1054])).
fof(f2824, plain, (~ spl10_102 | ~ spl10_120 | spl10_153), inference(avatar_split_clause, [], [f2820, f1344, f1144, f1069])).
fof(f2820, plain, (~ (e1 = op(e0, e4)) | (~ spl10_120 | spl10_153)), inference(backward_demodulation, [], [f1346, f1146])).
fof(f2814, plain, (~ spl10_8 | ~ spl10_129), inference(avatar_contradiction_clause, [], [f2813])).
fof(f2813, plain, ($false | (~ spl10_8 | ~ spl10_129)), inference(subsumption_resolution, [], [f2812, f462])).
fof(f2812, plain, ((e2 = e4) | (~ spl10_8 | ~ spl10_129)), inference(forward_demodulation, [], [f2793, f676])).
fof(f2793, plain, ((e4 = op(e4, e3)) | ~ spl10_129), inference(backward_demodulation, [], [f302, f1184])).
fof(f2811, plain, (~ spl10_28 | ~ spl10_129), inference(avatar_contradiction_clause, [], [f2810])).
fof(f2810, plain, ($false | (~ spl10_28 | ~ spl10_129)), inference(subsumption_resolution, [], [f2809, f462])).
fof(f2809, plain, ((e2 = e4) | (~ spl10_28 | ~ spl10_129)), inference(forward_demodulation, [], [f2792, f760])).
fof(f2792, plain, ((e4 = op(e3, e4)) | ~ spl10_129), inference(backward_demodulation, [], [f301, f1184])).
fof(f2804, plain, (~ spl10_57 | ~ spl10_129), inference(avatar_contradiction_clause, [], [f2803])).
fof(f2803, plain, ($false | (~ spl10_57 | ~ spl10_129)), inference(subsumption_resolution, [], [f2802, f458])).
fof(f2802, plain, ((e1 = e2) | (~ spl10_57 | ~ spl10_129)), inference(forward_demodulation, [], [f2789, f882])).
fof(f2789, plain, ((e2 = op(e2, e3)) | ~ spl10_129), inference(backward_demodulation, [], [f298, f1184])).
fof(f2762, plain, (spl10_22 | ~ spl10_76 | ~ spl10_132), inference(avatar_split_clause, [], [f2761, f1245, f960, f733])).
fof(f2761, plain, ((e1 = op(e4, e0)) | (~ spl10_76 | ~ spl10_132)), inference(forward_demodulation, [], [f1247, f962])).
fof(f2718, plain, (~ spl10_26 | ~ spl10_54 | ~ spl10_65 | spl10_215), inference(avatar_split_clause, [], [f2649, f1661, f913, f867, f750])).
fof(f2649, plain, (~ (e0 = op(e3, e4)) | (~ spl10_54 | ~ spl10_65 | spl10_215)), inference(forward_demodulation, [], [f2641, f869])).
fof(f2641, plain, (~ (e0 = op(op(e2, e4), e4)) | (~ spl10_65 | spl10_215)), inference(backward_demodulation, [], [f1663, f915])).
fof(f2707, plain, (~ spl10_16 | ~ spl10_2 | spl10_225), inference(avatar_split_clause, [], [f2689, f1711, f649, f708])).
fof(f2689, plain, (~ (e0 = op(e4, e1)) | (~ spl10_2 | spl10_225)), inference(backward_demodulation, [], [f1713, f651])).
fof(f2681, plain, (~ spl10_6 | ~ spl10_8), inference(avatar_contradiction_clause, [], [f2680])).
fof(f2680, plain, ($false | (~ spl10_6 | ~ spl10_8)), inference(subsumption_resolution, [], [f2679, f455])).
fof(f2679, plain, ((e0 = e2) | (~ spl10_6 | ~ spl10_8)), inference(forward_demodulation, [], [f676, f668])).
fof(f2678, plain, (~ spl10_11 | ~ spl10_14), inference(avatar_contradiction_clause, [], [f2677])).
fof(f2677, plain, ($false | (~ spl10_11 | ~ spl10_14)), inference(subsumption_resolution, [], [f2675, f456])).
fof(f2675, plain, ((e0 = e3) | (~ spl10_11 | ~ spl10_14)), inference(backward_demodulation, [], [f701, f689])).
fof(f2672, plain, (~ spl10_26 | ~ spl10_28), inference(avatar_contradiction_clause, [], [f2671])).
fof(f2671, plain, ($false | (~ spl10_26 | ~ spl10_28)), inference(subsumption_resolution, [], [f2670, f455])).
fof(f2670, plain, ((e0 = e2) | (~ spl10_26 | ~ spl10_28)), inference(forward_demodulation, [], [f760, f752])).
fof(f2659, plain, (~ spl10_42 | ~ spl10_45), inference(avatar_contradiction_clause, [], [f2658])).
fof(f2658, plain, ($false | (~ spl10_42 | ~ spl10_45)), inference(subsumption_resolution, [], [f2657, f460])).
fof(f2657, plain, ((e1 = e4) | (~ spl10_42 | ~ spl10_45)), inference(backward_demodulation, [], [f831, f819])).
fof(f2636, plain, (~ spl10_76 | ~ spl10_79), inference(avatar_contradiction_clause, [], [f2635])).
fof(f2635, plain, ($false | (~ spl10_76 | ~ spl10_79)), inference(subsumption_resolution, [], [f2634, f456])).
fof(f2634, plain, ((e0 = e3) | (~ spl10_76 | ~ spl10_79)), inference(backward_demodulation, [], [f974, f962])).
fof(f2622, plain, (~ spl10_65 | ~ spl10_93 | spl10_242), inference(avatar_split_clause, [], [f2618, f1841, f1031, f913])).
fof(f2618, plain, (~ (e4 = op(e2, e2)) | (~ spl10_93 | spl10_242)), inference(backward_demodulation, [], [f1843, f1033])).
fof(f2621, plain, (~ spl10_64 | ~ spl10_93 | spl10_238), inference(avatar_split_clause, [], [f2617, f1801, f1031, f909])).
fof(f2617, plain, (~ (e3 = op(e2, e2)) | (~ spl10_93 | spl10_238)), inference(backward_demodulation, [], [f1803, f1033])).
fof(f2612, plain, (~ spl10_107 | ~ spl10_109), inference(avatar_contradiction_clause, [], [f2611])).
fof(f2611, plain, ($false | (~ spl10_107 | ~ spl10_109)), inference(subsumption_resolution, [], [f2610, f459])).
fof(f2610, plain, ((e1 = e3) | (~ spl10_107 | ~ spl10_109)), inference(backward_demodulation, [], [f1100, f1092])).
fof(f2607, plain, (spl10_86 | ~ spl10_118 | ~ spl10_152), inference(avatar_split_clause, [], [f2604, f1339, f1136, f1002])).
fof(f2604, plain, ((e0 = op(e1, e2)) | (~ spl10_118 | ~ spl10_152)), inference(backward_demodulation, [], [f1341, f1138])).
fof(f2601, plain, (~ spl10_19 | ~ spl10_127), inference(avatar_contradiction_clause, [], [f2600])).
fof(f2600, plain, ($false | (~ spl10_19 | ~ spl10_127)), inference(subsumption_resolution, [], [f2599, f463])).
fof(f2599, plain, ((e3 = e4) | (~ spl10_19 | ~ spl10_127)), inference(forward_demodulation, [], [f2583, f722])).
fof(f2583, plain, ((e4 = op(e4, e1)) | ~ spl10_127), inference(backward_demodulation, [], [f302, f1176])).
fof(f2598, plain, (~ spl10_79 | ~ spl10_127), inference(avatar_contradiction_clause, [], [f2597])).
fof(f2597, plain, ($false | (~ spl10_79 | ~ spl10_127)), inference(subsumption_resolution, [], [f2596, f463])).
fof(f2596, plain, ((e3 = e4) | (~ spl10_79 | ~ spl10_127)), inference(forward_demodulation, [], [f2582, f974])).
fof(f2582, plain, ((e4 = op(e1, e4)) | ~ spl10_127), inference(backward_demodulation, [], [f301, f1176])).
fof(f2593, plain, (~ spl10_66 | ~ spl10_127), inference(avatar_contradiction_clause, [], [f2592])).
fof(f2592, plain, ($false | (~ spl10_66 | ~ spl10_127)), inference(subsumption_resolution, [], [f2591, f455])).
fof(f2591, plain, ((e0 = e2) | (~ spl10_66 | ~ spl10_127)), inference(forward_demodulation, [], [f2579, f920])).
fof(f2579, plain, ((e2 = op(e2, e1)) | ~ spl10_127), inference(backward_demodulation, [], [f298, f1176])).
fof(f2573, plain, (~ spl10_118 | ~ spl10_113), inference(avatar_split_clause, [], [f2572, f1115, f1136])).
fof(f2572, plain, (~ (e2 = op(e0, e1)) | ~ spl10_113), inference(forward_demodulation, [], [f408, f1117])).
fof(f2546, plain, (~ spl10_6 | ~ spl10_79 | ~ spl10_132), inference(avatar_contradiction_clause, [], [f2545])).
fof(f2545, plain, ($false | (~ spl10_6 | ~ spl10_79 | ~ spl10_132)), inference(subsumption_resolution, [], [f2544, f454])).
fof(f2544, plain, ((e0 = e1) | (~ spl10_6 | ~ spl10_79 | ~ spl10_132)), inference(forward_demodulation, [], [f2539, f668])).
fof(f2539, plain, ((e1 = op(e4, e3)) | (~ spl10_79 | ~ spl10_132)), inference(backward_demodulation, [], [f1247, f974])).
fof(f2498, plain, (~ spl10_34 | ~ spl10_109), inference(avatar_split_clause, [], [f2497, f1098, f783])).
fof(f2497, plain, (~ (e3 = op(e3, e3)) | ~ spl10_109), inference(forward_demodulation, [], [f386, f1100])).
fof(f2491, plain, (~ spl10_26 | ~ spl10_85 | ~ spl10_138), inference(avatar_contradiction_clause, [], [f2490])).
fof(f2490, plain, ($false | (~ spl10_26 | ~ spl10_85 | ~ spl10_138)), inference(subsumption_resolution, [], [f2489, f454])).
fof(f2489, plain, ((e0 = e1) | (~ spl10_26 | ~ spl10_85 | ~ spl10_138)), inference(forward_demodulation, [], [f2402, f752])).
fof(f2402, plain, ((e1 = op(e3, e4)) | (~ spl10_85 | ~ spl10_138)), inference(backward_demodulation, [], [f1275, f999])).
fof(f2458, plain, (~ spl10_6 | ~ spl10_7), inference(avatar_contradiction_clause, [], [f2457])).
fof(f2457, plain, ($false | (~ spl10_6 | ~ spl10_7)), inference(subsumption_resolution, [], [f2456, f454])).
fof(f2456, plain, ((e0 = e1) | (~ spl10_6 | ~ spl10_7)), inference(backward_demodulation, [], [f672, f668])).
fof(f2448, plain, (~ spl10_26 | ~ spl10_27), inference(avatar_contradiction_clause, [], [f2447])).
fof(f2447, plain, ($false | (~ spl10_26 | ~ spl10_27)), inference(subsumption_resolution, [], [f2446, f454])).
fof(f2446, plain, ((e0 = e1) | (~ spl10_26 | ~ spl10_27)), inference(backward_demodulation, [], [f756, f752])).
fof(f2442, plain, (~ spl10_43 | ~ spl10_45), inference(avatar_contradiction_clause, [], [f2441])).
fof(f2441, plain, ($false | (~ spl10_43 | ~ spl10_45)), inference(subsumption_resolution, [], [f2440, f462])).
fof(f2440, plain, ((e2 = e4) | (~ spl10_43 | ~ spl10_45)), inference(forward_demodulation, [], [f831, f823])).
fof(f2419, plain, (~ spl10_72 | ~ spl10_73), inference(avatar_contradiction_clause, [], [f2418])).
fof(f2418, plain, ($false | (~ spl10_72 | ~ spl10_73)), inference(subsumption_resolution, [], [f2417, f458])).
fof(f2417, plain, ((e1 = e2) | (~ spl10_72 | ~ spl10_73)), inference(backward_demodulation, [], [f949, f945])).
fof(f2408, plain, (~ spl10_83 | ~ spl10_85), inference(avatar_contradiction_clause, [], [f2407])).
fof(f2407, plain, ($false | (~ spl10_83 | ~ spl10_85)), inference(subsumption_resolution, [], [f2406, f462])).
fof(f2406, plain, ((e2 = e4) | (~ spl10_83 | ~ spl10_85)), inference(backward_demodulation, [], [f999, f991])).
fof(f2389, plain, (~ spl10_82 | ~ spl10_97), inference(avatar_split_clause, [], [f2383, f1048, f985])).
fof(f2383, plain, (~ (e1 = op(e1, e3)) | ~ spl10_97), inference(backward_demodulation, [], [f416, f1050])).
fof(f2387, plain, (~ spl10_47 | ~ spl10_97), inference(avatar_split_clause, [], [f2381, f1048, f838])).
fof(f2381, plain, (~ (e1 = op(e3, e0)) | ~ spl10_97), inference(backward_demodulation, [], [f359, f1050])).
fof(f2368, plain, (~ spl10_117 | spl10_153), inference(avatar_contradiction_clause, [], [f2367])).
fof(f2367, plain, ($false | (~ spl10_117 | spl10_153)), inference(subsumption_resolution, [], [f2362, f1134])).
fof(f2362, plain, (~ (e1 = op(e0, e1)) | (~ spl10_117 | spl10_153)), inference(backward_demodulation, [], [f1346, f1134])).
fof(f2364, plain, (~ spl10_112 | ~ spl10_117), inference(avatar_split_clause, [], [f2359, f1132, f1111])).
fof(f2359, plain, (~ (e1 = op(e0, e2)) | ~ spl10_117), inference(backward_demodulation, [], [f408, f1134])).
fof(f2349, plain, (~ spl10_111 | ~ spl10_86), inference(avatar_split_clause, [], [f2348, f1002, f1107])).
fof(f2348, plain, (~ (e0 = op(e0, e2)) | ~ spl10_86), inference(forward_demodulation, [], [f374, f1004])).
fof(f2347, plain, (~ spl10_112 | ~ spl10_12), inference(avatar_split_clause, [], [f2346, f691, f1111])).
fof(f2346, plain, (~ (e1 = op(e0, e2)) | ~ spl10_12), inference(forward_demodulation, [], [f377, f693])).
fof(f2337, plain, (~ spl10_103 | ~ spl10_105), inference(avatar_contradiction_clause, [], [f2336])).
fof(f2336, plain, ($false | (~ spl10_103 | ~ spl10_105)), inference(subsumption_resolution, [], [f2335, f462])).
fof(f2335, plain, ((e2 = e4) | (~ spl10_103 | ~ spl10_105)), inference(forward_demodulation, [], [f1083, f1075])).
fof(f2332, plain, (~ spl10_96 | ~ spl10_86), inference(avatar_split_clause, [], [f2331, f1002, f1044])).
fof(f2331, plain, (~ (e0 = op(e1, e0)) | ~ spl10_86), inference(forward_demodulation, [], [f415, f1004])).
fof(f2319, plain, (~ spl10_92 | ~ spl10_77), inference(avatar_split_clause, [], [f2318, f964, f1027])).
fof(f2318, plain, (~ (e1 = op(e1, e1)) | ~ spl10_77), inference(forward_demodulation, [], [f420, f966])).
fof(f2315, plain, (~ spl10_82 | ~ spl10_7), inference(avatar_split_clause, [], [f2314, f670, f985])).
fof(f2314, plain, (~ (e1 = op(e1, e3)) | ~ spl10_7), inference(forward_demodulation, [], [f390, f672])).
fof(f2308, plain, (~ spl10_27 | ~ spl10_77), inference(avatar_contradiction_clause, [], [f2307])).
fof(f2307, plain, ($false | (~ spl10_27 | ~ spl10_77)), inference(subsumption_resolution, [], [f2306, f966])).
fof(f2306, plain, (~ (e1 = op(e1, e4)) | ~ spl10_27), inference(forward_demodulation, [], [f399, f756])).
fof(f2301, plain, (~ spl10_95 | ~ spl10_77 | spl10_133), inference(avatar_split_clause, [], [f2300, f1250, f964, f1039])).
fof(f2300, plain, (~ (e4 = op(e1, e1)) | (~ spl10_77 | spl10_133)), inference(forward_demodulation, [], [f1252, f966])).
fof(f2256, plain, (~ spl10_23 | ~ spl10_25), inference(avatar_contradiction_clause, [], [f2255])).
fof(f2255, plain, ($false | (~ spl10_23 | ~ spl10_25)), inference(subsumption_resolution, [], [f2254, f462])).
fof(f2254, plain, ((e2 = e4) | (~ spl10_23 | ~ spl10_25)), inference(forward_demodulation, [], [f747, f739])).
fof(f2253, plain, (~ spl10_72 | ~ spl10_1 | ~ spl10_23 | spl10_195), inference(avatar_split_clause, [], [f2252, f1561, f737, f645, f943])).
fof(f2252, plain, (~ (e1 = op(e2, e0)) | (~ spl10_1 | ~ spl10_23 | spl10_195)), inference(backward_demodulation, [], [f2227, f739])).
fof(f2227, plain, (~ (e1 = op(op(e4, e0), e0)) | (~ spl10_1 | spl10_195)), inference(backward_demodulation, [], [f1563, f647])).
fof(f2251, plain, (~ spl10_22 | ~ spl10_1 | spl10_205), inference(avatar_split_clause, [], [f2228, f1607, f645, f733])).
fof(f2228, plain, (~ (e1 = op(e4, e0)) | (~ spl10_1 | spl10_205)), inference(backward_demodulation, [], [f1609, f647])).
fof(f2248, plain, (~ spl10_17 | ~ spl10_7), inference(avatar_split_clause, [], [f2247, f670, f712])).
fof(f2247, plain, (~ (e1 = op(e4, e1)) | ~ spl10_7), inference(forward_demodulation, [], [f449, f672])).
fof(f2239, plain, (~ spl10_12 | ~ spl10_14), inference(avatar_contradiction_clause, [], [f2238])).
fof(f2238, plain, ($false | (~ spl10_12 | ~ spl10_14)), inference(subsumption_resolution, [], [f2237, f459])).
fof(f2237, plain, ((e1 = e3) | (~ spl10_12 | ~ spl10_14)), inference(forward_demodulation, [], [f701, f693])).
fof(f2225, plain, (~ spl10_7 | ~ spl10_8), inference(avatar_contradiction_clause, [], [f2224])).
fof(f2224, plain, ($false | (~ spl10_7 | ~ spl10_8)), inference(subsumption_resolution, [], [f2223, f458])).
fof(f2223, plain, ((e1 = e2) | (~ spl10_7 | ~ spl10_8)), inference(backward_demodulation, [], [f676, f672])).
fof(f2222, plain, (~ spl10_8 | ~ spl10_10), inference(avatar_contradiction_clause, [], [f2221])).
fof(f2221, plain, ($false | (~ spl10_8 | ~ spl10_10)), inference(subsumption_resolution, [], [f2220, f462])).
fof(f2220, plain, ((e2 = e4) | (~ spl10_8 | ~ spl10_10)), inference(backward_demodulation, [], [f684, f676])).
fof(f2217, plain, (~ spl10_12 | ~ spl10_13), inference(avatar_contradiction_clause, [], [f2216])).
fof(f2216, plain, ($false | (~ spl10_12 | ~ spl10_13)), inference(subsumption_resolution, [], [f2215, f458])).
fof(f2215, plain, ((e1 = e2) | (~ spl10_12 | ~ spl10_13)), inference(backward_demodulation, [], [f697, f693])).
fof(f2214, plain, (~ spl10_13 | ~ spl10_14), inference(avatar_contradiction_clause, [], [f2213])).
fof(f2213, plain, ($false | (~ spl10_13 | ~ spl10_14)), inference(subsumption_resolution, [], [f2212, f461])).
fof(f2212, plain, ((e2 = e3) | (~ spl10_13 | ~ spl10_14)), inference(backward_demodulation, [], [f701, f697])).
fof(f2211, plain, (~ spl10_14 | ~ spl10_15), inference(avatar_contradiction_clause, [], [f2210])).
fof(f2210, plain, ($false | (~ spl10_14 | ~ spl10_15)), inference(subsumption_resolution, [], [f2209, f463])).
fof(f2209, plain, ((e3 = e4) | (~ spl10_14 | ~ spl10_15)), inference(backward_demodulation, [], [f705, f701])).
fof(f2192, plain, (~ spl10_22 | ~ spl10_23), inference(avatar_contradiction_clause, [], [f2191])).
fof(f2191, plain, ($false | (~ spl10_22 | ~ spl10_23)), inference(subsumption_resolution, [], [f2190, f458])).
fof(f2190, plain, ((e1 = e2) | (~ spl10_22 | ~ spl10_23)), inference(backward_demodulation, [], [f739, f735])).
fof(f2189, plain, (~ spl10_23 | ~ spl10_24), inference(avatar_contradiction_clause, [], [f2188])).
fof(f2188, plain, ($false | (~ spl10_23 | ~ spl10_24)), inference(subsumption_resolution, [], [f2187, f461])).
fof(f2187, plain, ((e2 = e3) | (~ spl10_23 | ~ spl10_24)), inference(backward_demodulation, [], [f743, f739])).
fof(f2186, plain, (~ spl10_24 | ~ spl10_25), inference(avatar_contradiction_clause, [], [f2185])).
fof(f2185, plain, ($false | (~ spl10_24 | ~ spl10_25)), inference(subsumption_resolution, [], [f2184, f463])).
fof(f2184, plain, ((e3 = e4) | (~ spl10_24 | ~ spl10_25)), inference(backward_demodulation, [], [f747, f743])).
fof(f2175, plain, (~ spl10_27 | ~ spl10_30), inference(avatar_contradiction_clause, [], [f2174])).
fof(f2174, plain, ($false | (~ spl10_27 | ~ spl10_30)), inference(subsumption_resolution, [], [f2173, f460])).
fof(f2173, plain, ((e1 = e4) | (~ spl10_27 | ~ spl10_30)), inference(backward_demodulation, [], [f768, f756])).
fof(f2130, plain, (~ spl10_52 | ~ spl10_53), inference(avatar_contradiction_clause, [], [f2129])).
fof(f2129, plain, ($false | (~ spl10_52 | ~ spl10_53)), inference(subsumption_resolution, [], [f2128, f458])).
fof(f2128, plain, ((e1 = e2) | (~ spl10_52 | ~ spl10_53)), inference(backward_demodulation, [], [f865, f861])).
fof(f2124, plain, (~ spl10_54 | ~ spl10_55), inference(avatar_contradiction_clause, [], [f2123])).
fof(f2123, plain, ($false | (~ spl10_54 | ~ spl10_55)), inference(subsumption_resolution, [], [f2122, f463])).
fof(f2122, plain, ((e3 = e4) | (~ spl10_54 | ~ spl10_55)), inference(backward_demodulation, [], [f873, f869])).
fof(f2065, plain, (~ spl10_78 | ~ spl10_79), inference(avatar_contradiction_clause, [], [f2064])).
fof(f2064, plain, ($false | (~ spl10_78 | ~ spl10_79)), inference(subsumption_resolution, [], [f2063, f461])).
fof(f2063, plain, ((e2 = e3) | (~ spl10_78 | ~ spl10_79)), inference(backward_demodulation, [], [f974, f970])).
fof(f2059, plain, (~ spl10_80 | spl10_133), inference(avatar_contradiction_clause, [], [f2058])).
fof(f2058, plain, ($false | (~ spl10_80 | spl10_133)), inference(subsumption_resolution, [], [f2053, f978])).
fof(f2053, plain, (~ (e4 = op(e1, e4)) | (~ spl10_80 | spl10_133)), inference(backward_demodulation, [], [f1252, f978])).
fof(f2030, plain, (~ spl10_61 | ~ spl10_86), inference(avatar_split_clause, [], [f2023, f1002, f897])).
fof(f2023, plain, (~ (e0 = op(e2, e2)) | ~ spl10_86), inference(backward_demodulation, [], [f378, f1004])).
fof(f2019, plain, (~ spl10_91 | ~ spl10_121 | spl10_198), inference(avatar_contradiction_clause, [], [f2018])).
fof(f2018, plain, ($false | (~ spl10_91 | ~ spl10_121 | spl10_198)), inference(subsumption_resolution, [], [f2008, f1151])).
fof(f2008, plain, (~ (e0 = op(e0, e0)) | (~ spl10_91 | spl10_198)), inference(backward_demodulation, [], [f1577, f1025])).
fof(f2001, plain, (~ spl10_96 | spl10_156), inference(avatar_contradiction_clause, [], [f2000])).
fof(f2000, plain, ($false | (~ spl10_96 | spl10_156)), inference(subsumption_resolution, [], [f1989, f1046])).
fof(f1989, plain, (~ (e0 = op(e1, e0)) | (~ spl10_96 | spl10_156)), inference(backward_demodulation, [], [f1360, f1046])).
fof(f1992, plain, (~ spl10_21 | ~ spl10_96), inference(avatar_split_clause, [], [f1983, f1044, f729])).
fof(f1983, plain, (~ (e0 = op(e4, e0)) | ~ spl10_96), inference(backward_demodulation, [], [f360, f1046])).
fof(f1980, plain, (~ spl10_102 | ~ spl10_103), inference(avatar_contradiction_clause, [], [f1979])).
fof(f1979, plain, ($false | (~ spl10_102 | ~ spl10_103)), inference(subsumption_resolution, [], [f1978, f458])).
fof(f1978, plain, ((e1 = e2) | (~ spl10_102 | ~ spl10_103)), inference(backward_demodulation, [], [f1075, f1071])).
fof(f1977, plain, (~ spl10_103 | ~ spl10_104), inference(avatar_contradiction_clause, [], [f1976])).
fof(f1976, plain, ($false | (~ spl10_103 | ~ spl10_104)), inference(subsumption_resolution, [], [f1975, f461])).
fof(f1975, plain, ((e2 = e3) | (~ spl10_103 | ~ spl10_104)), inference(backward_demodulation, [], [f1079, f1075])).
fof(f1974, plain, (~ spl10_104 | ~ spl10_105), inference(avatar_contradiction_clause, [], [f1973])).
fof(f1973, plain, ($false | (~ spl10_104 | ~ spl10_105)), inference(subsumption_resolution, [], [f1972, f463])).
fof(f1972, plain, ((e3 = e4) | (~ spl10_104 | ~ spl10_105)), inference(backward_demodulation, [], [f1083, f1079])).
fof(f1959, plain, (~ spl10_81 | ~ spl10_106), inference(avatar_split_clause, [], [f1954, f1086, f981])).
fof(f1954, plain, (~ (e0 = op(e1, e3)) | ~ spl10_106), inference(backward_demodulation, [], [f384, f1088])).
fof(f1939, plain, (~ spl10_101 | ~ spl10_116), inference(avatar_split_clause, [], [f1930, f1128, f1065])).
fof(f1930, plain, (~ (e0 = op(e0, e4)) | ~ spl10_116), inference(backward_demodulation, [], [f410, f1130])).
fof(f1938, plain, (~ spl10_106 | ~ spl10_116), inference(avatar_split_clause, [], [f1929, f1128, f1086])).
fof(f1929, plain, (~ (e0 = op(e0, e3)) | ~ spl10_116), inference(backward_demodulation, [], [f409, f1130])).
fof(f1935, plain, (~ spl10_41 | ~ spl10_116), inference(avatar_split_clause, [], [f1926, f1128, f813])).
fof(f1926, plain, (~ (e0 = op(e3, e1)) | ~ spl10_116), inference(backward_demodulation, [], [f366, f1130])).
fof(f1918, plain, (~ spl10_106 | ~ spl10_121), inference(avatar_split_clause, [], [f1906, f1149, f1086])).
fof(f1906, plain, (~ (e0 = op(e0, e3)) | ~ spl10_121), inference(backward_demodulation, [], [f406, f1151])).
fof(f1916, plain, (~ spl10_116 | ~ spl10_121), inference(avatar_split_clause, [], [f1904, f1149, f1128])).
fof(f1904, plain, (~ (e0 = op(e0, e1)) | ~ spl10_121), inference(backward_demodulation, [], [f404, f1151])).
fof(f1914, plain, (~ spl10_46 | ~ spl10_121), inference(avatar_split_clause, [], [f1902, f1149, f834])).
fof(f1902, plain, (~ (e0 = op(e3, e0)) | ~ spl10_121), inference(backward_demodulation, [], [f356, f1151])).
fof(f1912, plain, (~ spl10_96 | ~ spl10_121), inference(avatar_split_clause, [], [f1900, f1149, f1044])).
fof(f1900, plain, (~ (e0 = op(e1, e0)) | ~ spl10_121), inference(backward_demodulation, [], [f354, f1151])).
fof(f1875, plain, (~ spl10_245 | ~ spl10_226 | ~ spl10_227 | ~ spl10_122), inference(avatar_split_clause, [], [f639, f1153, f1721, f1717, f1859])).
fof(f639, plain, (~ (op(e0, e0) = e1) | ~ (e2 = op(e0, op(e0, e0))) | ~ (e3 = op(op(e0, op(e0, e0)), op(e0, e0))) | ~ (e4 = op(op(e0, e0), op(e0, e0)))), inference(cnf_transformation, [], [f242])).
fof(f242, plain, (~ (op(e0, e0) = e1) | ~ (e2 = op(e0, op(e0, e0))) | ~ (e3 = op(op(e0, op(e0, e0)), op(e0, e0))) | ~ (e4 = op(op(e0, e0), op(e0, e0)))), inference(ennf_transformation, [], [f122])).
fof(f122, plain, ~ ((op(e0, e0) = e1) & (e2 = op(e0, op(e0, e0))) & (e3 = op(op(e0, op(e0, e0)), op(e0, e0))) & (e4 = op(op(e0, e0), op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax122)).
fof(f1874, plain, (~ spl10_245 | ~ spl10_226 | ~ spl10_236 | ~ spl10_123), inference(avatar_split_clause, [], [f638, f1157, f1781, f1717, f1859])).
fof(f638, plain, (~ (op(e0, e0) = e2) | ~ (e1 = op(e0, op(e0, e0))) | ~ (e3 = op(op(e0, op(e0, e0)), op(e0, e0))) | ~ (e4 = op(op(e0, e0), op(e0, e0)))), inference(cnf_transformation, [], [f241])).
fof(f241, plain, (~ (op(e0, e0) = e2) | ~ (e1 = op(e0, op(e0, e0))) | ~ (e3 = op(op(e0, op(e0, e0)), op(e0, e0))) | ~ (e4 = op(op(e0, e0), op(e0, e0)))), inference(ennf_transformation, [], [f121])).
fof(f121, plain, ~ ((op(e0, e0) = e2) & (e1 = op(e0, op(e0, e0))) & (e3 = op(op(e0, op(e0, e0)), op(e0, e0))) & (e4 = op(op(e0, e0), op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax121)).
fof(f1873, plain, (~ spl10_244 | ~ spl10_202 | ~ spl10_203 | ~ spl10_31), inference(avatar_split_clause, [], [f637, f771, f1597, f1593, f1853])).
fof(f637, plain, (~ (e0 = op(e3, e3)) | ~ (e1 = op(e3, op(e3, e3))) | ~ (e2 = op(op(e3, op(e3, e3)), op(e3, e3))) | ~ (e4 = op(op(e3, e3), op(e3, e3)))), inference(cnf_transformation, [], [f240])).
fof(f240, plain, (~ (e0 = op(e3, e3)) | ~ (e1 = op(e3, op(e3, e3))) | ~ (e2 = op(op(e3, op(e3, e3)), op(e3, e3))) | ~ (e4 = op(op(e3, e3), op(e3, e3)))), inference(ennf_transformation, [], [f120])).
fof(f120, plain, ~ ((e0 = op(e3, e3)) & (e1 = op(e3, op(e3, e3))) & (e2 = op(op(e3, op(e3, e3)), op(e3, e3))) & (e4 = op(op(e3, e3), op(e3, e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax120)).
fof(f1872, plain, (~ spl10_244 | ~ spl10_202 | ~ spl10_224 | ~ spl10_32), inference(avatar_split_clause, [], [f636, f775, f1705, f1593, f1853])).
fof(f636, plain, (~ (e1 = op(e3, e3)) | ~ (e0 = op(e3, op(e3, e3))) | ~ (e2 = op(op(e3, op(e3, e3)), op(e3, e3))) | ~ (e4 = op(op(e3, e3), op(e3, e3)))), inference(cnf_transformation, [], [f239])).
fof(f239, plain, (~ (e1 = op(e3, e3)) | ~ (e0 = op(e3, op(e3, e3))) | ~ (e2 = op(op(e3, op(e3, e3)), op(e3, e3))) | ~ (e4 = op(op(e3, e3), op(e3, e3)))), inference(ennf_transformation, [], [f119])).
fof(f119, plain, ~ ((e1 = op(e3, e3)) & (e0 = op(e3, op(e3, e3))) & (e2 = op(op(e3, op(e3, e3)), op(e3, e3))) & (e4 = op(op(e3, e3), op(e3, e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax119)).
fof(f1871, plain, (~ spl10_242 | ~ spl10_199 | ~ spl10_200 | ~ spl10_91), inference(avatar_split_clause, [], [f635, f1023, f1583, f1579, f1841])).
fof(f635, plain, (~ (e0 = op(e1, e1)) | ~ (e3 = op(e1, op(e1, e1))) | ~ (e2 = op(op(e1, op(e1, e1)), op(e1, e1))) | ~ (e4 = op(op(e1, e1), op(e1, e1)))), inference(cnf_transformation, [], [f238])).
fof(f238, plain, (~ (e0 = op(e1, e1)) | ~ (e3 = op(e1, op(e1, e1))) | ~ (e2 = op(op(e1, op(e1, e1)), op(e1, e1))) | ~ (e4 = op(op(e1, e1), op(e1, e1)))), inference(ennf_transformation, [], [f118])).
fof(f118, plain, ~ ((e0 = op(e1, e1)) & (e3 = op(e1, op(e1, e1))) & (e2 = op(op(e1, op(e1, e1)), op(e1, e1))) & (e4 = op(op(e1, e1), op(e1, e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax118)).
fof(f1870, plain, (~ spl10_242 | ~ spl10_199 | ~ spl10_237 | ~ spl10_94), inference(avatar_split_clause, [], [f634, f1035, f1787, f1579, f1841])).
fof(f634, plain, (~ (e3 = op(e1, e1)) | ~ (e0 = op(e1, op(e1, e1))) | ~ (e2 = op(op(e1, op(e1, e1)), op(e1, e1))) | ~ (e4 = op(op(e1, e1), op(e1, e1)))), inference(cnf_transformation, [], [f237])).
fof(f237, plain, (~ (e3 = op(e1, e1)) | ~ (e0 = op(e1, op(e1, e1))) | ~ (e2 = op(op(e1, op(e1, e1)), op(e1, e1))) | ~ (e4 = op(op(e1, e1), op(e1, e1)))), inference(ennf_transformation, [], [f117])).
fof(f117, plain, ~ ((e3 = op(e1, e1)) & (e0 = op(e1, op(e1, e1))) & (e2 = op(op(e1, op(e1, e1)), op(e1, e1))) & (e4 = op(op(e1, e1), op(e1, e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax117)).
fof(f1868, plain, (~ spl10_245 | ~ spl10_221 | ~ spl10_236 | ~ spl10_124), inference(avatar_split_clause, [], [f632, f1161, f1781, f1691, f1859])).
fof(f632, plain, (~ (op(e0, e0) = e3) | ~ (e1 = op(e0, op(e0, e0))) | ~ (e2 = op(op(e0, op(e0, e0)), op(e0, e0))) | ~ (e4 = op(op(e0, e0), op(e0, e0)))), inference(cnf_transformation, [], [f235])).
fof(f235, plain, (~ (op(e0, e0) = e3) | ~ (e1 = op(e0, op(e0, e0))) | ~ (e2 = op(op(e0, op(e0, e0)), op(e0, e0))) | ~ (e4 = op(op(e0, e0), op(e0, e0)))), inference(ennf_transformation, [], [f115])).
fof(f115, plain, ~ ((op(e0, e0) = e3) & (e1 = op(e0, op(e0, e0))) & (e2 = op(op(e0, op(e0, e0)), op(e0, e0))) & (e4 = op(op(e0, e0), op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax115)).
fof(f1867, plain, (~ spl10_244 | ~ spl10_191 | ~ spl10_192 | ~ spl10_31), inference(avatar_split_clause, [], [f631, f771, f1547, f1543, f1853])).
fof(f631, plain, (~ (e0 = op(e3, e3)) | ~ (e2 = op(e3, op(e3, e3))) | ~ (e1 = op(op(e3, op(e3, e3)), op(e3, e3))) | ~ (e4 = op(op(e3, e3), op(e3, e3)))), inference(cnf_transformation, [], [f234])).
fof(f234, plain, (~ (e0 = op(e3, e3)) | ~ (e2 = op(e3, op(e3, e3))) | ~ (e1 = op(op(e3, op(e3, e3)), op(e3, e3))) | ~ (e4 = op(op(e3, e3), op(e3, e3)))), inference(ennf_transformation, [], [f114])).
fof(f114, plain, ~ ((e0 = op(e3, e3)) & (e2 = op(e3, op(e3, e3))) & (e1 = op(op(e3, op(e3, e3)), op(e3, e3))) & (e4 = op(op(e3, e3), op(e3, e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax114)).
fof(f1864, plain, (~ spl10_243 | ~ spl10_187 | ~ spl10_228 | ~ spl10_64), inference(avatar_split_clause, [], [f628, f909, f1727, f1525, f1847])).
fof(f628, plain, (~ (e3 = op(e2, e2)) | ~ (e0 = op(e2, op(e2, e2))) | ~ (e1 = op(op(e2, op(e2, e2)), op(e2, e2))) | ~ (e4 = op(op(e2, e2), op(e2, e2)))), inference(cnf_transformation, [], [f231])).
fof(f231, plain, (~ (e3 = op(e2, e2)) | ~ (e0 = op(e2, op(e2, e2))) | ~ (e1 = op(op(e2, op(e2, e2)), op(e2, e2))) | ~ (e4 = op(op(e2, e2), op(e2, e2)))), inference(ennf_transformation, [], [f111])).
fof(f111, plain, ~ ((e3 = op(e2, e2)) & (e0 = op(e2, op(e2, e2))) & (e1 = op(op(e2, op(e2, e2)), op(e2, e2))) & (e4 = op(op(e2, e2), op(e2, e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax111)).
fof(f1863, plain, (~ spl10_245 | ~ spl10_235 | ~ spl10_222 | ~ spl10_123), inference(avatar_split_clause, [], [f627, f1157, f1695, f1771, f1859])).
fof(f627, plain, (~ (op(e0, e0) = e2) | ~ (e3 = op(e0, op(e0, e0))) | ~ (e1 = op(op(e0, op(e0, e0)), op(e0, e0))) | ~ (e4 = op(op(e0, e0), op(e0, e0)))), inference(cnf_transformation, [], [f230])).
fof(f230, plain, (~ (op(e0, e0) = e2) | ~ (e3 = op(e0, op(e0, e0))) | ~ (e1 = op(op(e0, op(e0, e0)), op(e0, e0))) | ~ (e4 = op(op(e0, e0), op(e0, e0)))), inference(ennf_transformation, [], [f110])).
fof(f110, plain, ~ ((op(e0, e0) = e2) & (e3 = op(e0, op(e0, e0))) & (e1 = op(op(e0, op(e0, e0)), op(e0, e0))) & (e4 = op(op(e0, e0), op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax110)).
fof(f1862, plain, (~ spl10_245 | ~ spl10_235 | ~ spl10_227 | ~ spl10_124), inference(avatar_split_clause, [], [f626, f1161, f1721, f1771, f1859])).
fof(f626, plain, (~ (op(e0, e0) = e3) | ~ (e2 = op(e0, op(e0, e0))) | ~ (e1 = op(op(e0, op(e0, e0)), op(e0, e0))) | ~ (e4 = op(op(e0, e0), op(e0, e0)))), inference(cnf_transformation, [], [f229])).
fof(f229, plain, (~ (op(e0, e0) = e3) | ~ (e2 = op(e0, op(e0, e0))) | ~ (e1 = op(op(e0, op(e0, e0)), op(e0, e0))) | ~ (e4 = op(op(e0, e0), op(e0, e0)))), inference(ennf_transformation, [], [f109])).
fof(f109, plain, ~ ((op(e0, e0) = e3) & (e2 = op(e0, op(e0, e0))) & (e1 = op(op(e0, op(e0, e0)), op(e0, e0))) & (e4 = op(op(e0, e0), op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax109)).
fof(f1857, plain, (~ spl10_244 | ~ spl10_217 | ~ spl10_192 | ~ spl10_32), inference(avatar_split_clause, [], [f625, f775, f1547, f1671, f1853])).
fof(f625, plain, (~ (e1 = op(e3, e3)) | ~ (e2 = op(e3, op(e3, e3))) | ~ (e0 = op(op(e3, op(e3, e3)), op(e3, e3))) | ~ (e4 = op(op(e3, e3), op(e3, e3)))), inference(cnf_transformation, [], [f228])).
fof(f228, plain, (~ (e1 = op(e3, e3)) | ~ (e2 = op(e3, op(e3, e3))) | ~ (e0 = op(op(e3, op(e3, e3)), op(e3, e3))) | ~ (e4 = op(op(e3, e3), op(e3, e3)))), inference(ennf_transformation, [], [f108])).
fof(f108, plain, ~ ((e1 = op(e3, e3)) & (e2 = op(e3, op(e3, e3))) & (e0 = op(op(e3, op(e3, e3)), op(e3, e3))) & (e4 = op(op(e3, e3), op(e3, e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax108)).
fof(f1856, plain, (~ spl10_244 | ~ spl10_217 | ~ spl10_203 | ~ spl10_33), inference(avatar_split_clause, [], [f624, f779, f1597, f1671, f1853])).
fof(f624, plain, (~ (e2 = op(e3, e3)) | ~ (e1 = op(e3, op(e3, e3))) | ~ (e0 = op(op(e3, op(e3, e3)), op(e3, e3))) | ~ (e4 = op(op(e3, e3), op(e3, e3)))), inference(cnf_transformation, [], [f227])).
fof(f227, plain, (~ (e2 = op(e3, e3)) | ~ (e1 = op(e3, op(e3, e3))) | ~ (e0 = op(op(e3, op(e3, e3)), op(e3, e3))) | ~ (e4 = op(op(e3, e3), op(e3, e3)))), inference(ennf_transformation, [], [f107])).
fof(f107, plain, ~ ((e2 = op(e3, e3)) & (e1 = op(e3, op(e3, e3))) & (e0 = op(op(e3, op(e3, e3)), op(e3, e3))) & (e4 = op(op(e3, e3), op(e3, e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax107)).
fof(f1851, plain, (~ spl10_243 | ~ spl10_215 | ~ spl10_188 | ~ spl10_62), inference(avatar_split_clause, [], [f623, f901, f1529, f1661, f1847])).
fof(f623, plain, (~ (e1 = op(e2, e2)) | ~ (e3 = op(e2, op(e2, e2))) | ~ (e0 = op(op(e2, op(e2, e2)), op(e2, e2))) | ~ (e4 = op(op(e2, e2), op(e2, e2)))), inference(cnf_transformation, [], [f226])).
fof(f226, plain, (~ (e1 = op(e2, e2)) | ~ (e3 = op(e2, op(e2, e2))) | ~ (e0 = op(op(e2, op(e2, e2)), op(e2, e2))) | ~ (e4 = op(op(e2, e2), op(e2, e2)))), inference(ennf_transformation, [], [f106])).
fof(f106, plain, ~ ((e1 = op(e2, e2)) & (e3 = op(e2, op(e2, e2))) & (e0 = op(op(e2, op(e2, e2)), op(e2, e2))) & (e4 = op(op(e2, e2), op(e2, e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax106)).
fof(f1850, plain, (~ spl10_243 | ~ spl10_215 | ~ spl10_209 | ~ spl10_64), inference(avatar_split_clause, [], [f622, f909, f1627, f1661, f1847])).
fof(f622, plain, (~ (e3 = op(e2, e2)) | ~ (e1 = op(e2, op(e2, e2))) | ~ (e0 = op(op(e2, op(e2, e2)), op(e2, e2))) | ~ (e4 = op(op(e2, e2), op(e2, e2)))), inference(cnf_transformation, [], [f225])).
fof(f225, plain, (~ (e3 = op(e2, e2)) | ~ (e1 = op(e2, op(e2, e2))) | ~ (e0 = op(op(e2, op(e2, e2)), op(e2, e2))) | ~ (e4 = op(op(e2, e2), op(e2, e2)))), inference(ennf_transformation, [], [f105])).
fof(f105, plain, ~ ((e3 = op(e2, e2)) & (e1 = op(e2, op(e2, e2))) & (e0 = op(op(e2, op(e2, e2)), op(e2, e2))) & (e4 = op(op(e2, e2), op(e2, e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax105)).
fof(f1845, plain, (~ spl10_242 | ~ spl10_231 | ~ spl10_200 | ~ spl10_93), inference(avatar_split_clause, [], [f621, f1031, f1583, f1749, f1841])).
fof(f621, plain, (~ (e2 = op(e1, e1)) | ~ (e3 = op(e1, op(e1, e1))) | ~ (e0 = op(op(e1, op(e1, e1)), op(e1, e1))) | ~ (e4 = op(op(e1, e1), op(e1, e1)))), inference(cnf_transformation, [], [f224])).
fof(f224, plain, (~ (e2 = op(e1, e1)) | ~ (e3 = op(e1, op(e1, e1))) | ~ (e0 = op(op(e1, op(e1, e1)), op(e1, e1))) | ~ (e4 = op(op(e1, e1), op(e1, e1)))), inference(ennf_transformation, [], [f104])).
fof(f104, plain, ~ ((e2 = op(e1, e1)) & (e3 = op(e1, op(e1, e1))) & (e0 = op(op(e1, op(e1, e1)), op(e1, e1))) & (e4 = op(op(e1, e1), op(e1, e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax104)).
fof(f1844, plain, (~ spl10_242 | ~ spl10_231 | ~ spl10_207 | ~ spl10_94), inference(avatar_split_clause, [], [f620, f1035, f1617, f1749, f1841])).
fof(f620, plain, (~ (e3 = op(e1, e1)) | ~ (e2 = op(e1, op(e1, e1))) | ~ (e0 = op(op(e1, op(e1, e1)), op(e1, e1))) | ~ (e4 = op(op(e1, e1), op(e1, e1)))), inference(cnf_transformation, [], [f223])).
fof(f223, plain, (~ (e3 = op(e1, e1)) | ~ (e2 = op(e1, op(e1, e1))) | ~ (e0 = op(op(e1, op(e1, e1)), op(e1, e1))) | ~ (e4 = op(op(e1, e1), op(e1, e1)))), inference(ennf_transformation, [], [f103])).
fof(f103, plain, ~ ((e3 = op(e1, e1)) & (e2 = op(e1, op(e1, e1))) & (e0 = op(op(e1, op(e1, e1)), op(e1, e1))) & (e4 = op(op(e1, e1), op(e1, e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax103)).
fof(f1838, plain, (~ spl10_239 | ~ spl10_212 | ~ spl10_228 | ~ spl10_62), inference(avatar_split_clause, [], [f618, f901, f1727, f1645, f1807])).
fof(f618, plain, (~ (e1 = op(e2, e2)) | ~ (e0 = op(e2, op(e2, e2))) | ~ (e4 = op(op(e2, op(e2, e2)), op(e2, e2))) | ~ (e3 = op(op(e2, e2), op(e2, e2)))), inference(cnf_transformation, [], [f221])).
fof(f221, plain, (~ (e1 = op(e2, e2)) | ~ (e0 = op(e2, op(e2, e2))) | ~ (e4 = op(op(e2, op(e2, e2)), op(e2, e2))) | ~ (e3 = op(op(e2, e2), op(e2, e2)))), inference(ennf_transformation, [], [f101])).
fof(f101, plain, ~ ((e1 = op(e2, e2)) & (e0 = op(e2, op(e2, e2))) & (e4 = op(op(e2, op(e2, e2)), op(e2, e2))) & (e3 = op(op(e2, e2), op(e2, e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax101)).
fof(f1837, plain, (~ spl10_238 | ~ spl10_211 | ~ spl10_207 | ~ spl10_91), inference(avatar_split_clause, [], [f617, f1023, f1617, f1639, f1801])).
fof(f617, plain, (~ (e0 = op(e1, e1)) | ~ (e2 = op(e1, op(e1, e1))) | ~ (e4 = op(op(e1, op(e1, e1)), op(e1, e1))) | ~ (e3 = op(op(e1, e1), op(e1, e1)))), inference(cnf_transformation, [], [f220])).
fof(f220, plain, (~ (e0 = op(e1, e1)) | ~ (e2 = op(e1, op(e1, e1))) | ~ (e4 = op(op(e1, op(e1, e1)), op(e1, e1))) | ~ (e3 = op(op(e1, e1), op(e1, e1)))), inference(ennf_transformation, [], [f100])).
fof(f100, plain, ~ ((e0 = op(e1, e1)) & (e2 = op(e1, op(e1, e1))) & (e4 = op(op(e1, op(e1, e1)), op(e1, e1))) & (e3 = op(op(e1, e1), op(e1, e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax100)).
fof(f1835, plain, (~ spl10_241 | ~ spl10_229 | ~ spl10_227 | ~ spl10_122), inference(avatar_split_clause, [], [f615, f1153, f1721, f1735, f1819])).
fof(f615, plain, (~ (op(e0, e0) = e1) | ~ (e2 = op(e0, op(e0, e0))) | ~ (e4 = op(op(e0, op(e0, e0)), op(e0, e0))) | ~ (e3 = op(op(e0, e0), op(e0, e0)))), inference(cnf_transformation, [], [f218])).
fof(f218, plain, (~ (op(e0, e0) = e1) | ~ (e2 = op(e0, op(e0, e0))) | ~ (e4 = op(op(e0, op(e0, e0)), op(e0, e0))) | ~ (e3 = op(op(e0, e0), op(e0, e0)))), inference(ennf_transformation, [], [f98])).
fof(f98, plain, ~ ((op(e0, e0) = e1) & (e2 = op(e0, op(e0, e0))) & (e4 = op(op(e0, op(e0, e0)), op(e0, e0))) & (e3 = op(op(e0, e0), op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax98)).
fof(f1832, plain, (~ spl10_240 | ~ spl10_204 | ~ spl10_225 | ~ spl10_2), inference(avatar_split_clause, [], [f612, f649, f1711, f1603, f1813])).
fof(f612, plain, (~ (e1 = op(e4, e4)) | ~ (e0 = op(e4, op(e4, e4))) | ~ (e2 = op(op(e4, op(e4, e4)), op(e4, e4))) | ~ (e3 = op(op(e4, e4), op(e4, e4)))), inference(cnf_transformation, [], [f215])).
fof(f215, plain, (~ (e1 = op(e4, e4)) | ~ (e0 = op(e4, op(e4, e4))) | ~ (e2 = op(op(e4, op(e4, e4)), op(e4, e4))) | ~ (e3 = op(op(e4, e4), op(e4, e4)))), inference(ennf_transformation, [], [f95])).
fof(f95, plain, ~ ((e1 = op(e4, e4)) & (e0 = op(e4, op(e4, e4))) & (e2 = op(op(e4, op(e4, e4)), op(e4, e4))) & (e3 = op(op(e4, e4), op(e4, e4)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax95)).
fof(f1831, plain, (~ spl10_238 | ~ spl10_199 | ~ spl10_201 | ~ spl10_91), inference(avatar_split_clause, [], [f611, f1023, f1588, f1579, f1801])).
fof(f611, plain, (~ (e0 = op(e1, e1)) | ~ (e4 = op(e1, op(e1, e1))) | ~ (e2 = op(op(e1, op(e1, e1)), op(e1, e1))) | ~ (e3 = op(op(e1, e1), op(e1, e1)))), inference(cnf_transformation, [], [f214])).
fof(f214, plain, (~ (e0 = op(e1, e1)) | ~ (e4 = op(e1, op(e1, e1))) | ~ (e2 = op(op(e1, op(e1, e1)), op(e1, e1))) | ~ (e3 = op(op(e1, e1), op(e1, e1)))), inference(ennf_transformation, [], [f94])).
fof(f94, plain, ~ ((e0 = op(e1, e1)) & (e4 = op(e1, op(e1, e1))) & (e2 = op(op(e1, op(e1, e1)), op(e1, e1))) & (e3 = op(op(e1, e1), op(e1, e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax94)).
fof(f1830, plain, (~ spl10_238 | ~ spl10_199 | ~ spl10_237 | ~ spl10_95), inference(avatar_split_clause, [], [f610, f1039, f1787, f1579, f1801])).
fof(f610, plain, (~ (e4 = op(e1, e1)) | ~ (e0 = op(e1, op(e1, e1))) | ~ (e2 = op(op(e1, op(e1, e1)), op(e1, e1))) | ~ (e3 = op(op(e1, e1), op(e1, e1)))), inference(cnf_transformation, [], [f213])).
fof(f213, plain, (~ (e4 = op(e1, e1)) | ~ (e0 = op(e1, op(e1, e1))) | ~ (e2 = op(op(e1, op(e1, e1)), op(e1, e1))) | ~ (e3 = op(op(e1, e1), op(e1, e1)))), inference(ennf_transformation, [], [f93])).
fof(f93, plain, ~ ((e4 = op(e1, e1)) & (e0 = op(e1, op(e1, e1))) & (e2 = op(op(e1, op(e1, e1)), op(e1, e1))) & (e3 = op(op(e1, e1), op(e1, e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax93)).
fof(f1829, plain, (~ spl10_241 | ~ spl10_221 | ~ spl10_223 | ~ spl10_122), inference(avatar_split_clause, [], [f609, f1153, f1700, f1691, f1819])).
fof(f609, plain, (~ (op(e0, e0) = e1) | ~ (e4 = op(e0, op(e0, e0))) | ~ (e2 = op(op(e0, op(e0, e0)), op(e0, e0))) | ~ (e3 = op(op(e0, e0), op(e0, e0)))), inference(cnf_transformation, [], [f212])).
fof(f212, plain, (~ (op(e0, e0) = e1) | ~ (e4 = op(e0, op(e0, e0))) | ~ (e2 = op(op(e0, op(e0, e0)), op(e0, e0))) | ~ (e3 = op(op(e0, e0), op(e0, e0)))), inference(ennf_transformation, [], [f92])).
fof(f92, plain, ~ ((op(e0, e0) = e1) & (e4 = op(e0, op(e0, e0))) & (e2 = op(op(e0, op(e0, e0)), op(e0, e0))) & (e3 = op(op(e0, e0), op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax92)).
fof(f1827, plain, (~ spl10_240 | ~ spl10_195 | ~ spl10_196 | ~ spl10_1), inference(avatar_split_clause, [], [f607, f645, f1565, f1561, f1813])).
fof(f607, plain, (~ (e0 = op(e4, e4)) | ~ (e2 = op(e4, op(e4, e4))) | ~ (e1 = op(op(e4, op(e4, e4)), op(e4, e4))) | ~ (e3 = op(op(e4, e4), op(e4, e4)))), inference(cnf_transformation, [], [f210])).
fof(f210, plain, (~ (e0 = op(e4, e4)) | ~ (e2 = op(e4, op(e4, e4))) | ~ (e1 = op(op(e4, op(e4, e4)), op(e4, e4))) | ~ (e3 = op(op(e4, e4), op(e4, e4)))), inference(ennf_transformation, [], [f90])).
fof(f90, plain, ~ ((e0 = op(e4, e4)) & (e2 = op(e4, op(e4, e4))) & (e1 = op(op(e4, op(e4, e4)), op(e4, e4))) & (e3 = op(op(e4, e4), op(e4, e4)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax90)).
fof(f1825, plain, (~ spl10_239 | ~ spl10_187 | ~ spl10_189 | ~ spl10_61), inference(avatar_split_clause, [], [f605, f897, f1534, f1525, f1807])).
fof(f605, plain, (~ (e0 = op(e2, e2)) | ~ (e4 = op(e2, op(e2, e2))) | ~ (e1 = op(op(e2, op(e2, e2)), op(e2, e2))) | ~ (e3 = op(op(e2, e2), op(e2, e2)))), inference(cnf_transformation, [], [f208])).
fof(f208, plain, (~ (e0 = op(e2, e2)) | ~ (e4 = op(e2, op(e2, e2))) | ~ (e1 = op(op(e2, op(e2, e2)), op(e2, e2))) | ~ (e3 = op(op(e2, e2), op(e2, e2)))), inference(ennf_transformation, [], [f88])).
fof(f88, plain, ~ ((e0 = op(e2, e2)) & (e4 = op(e2, op(e2, e2))) & (e1 = op(op(e2, op(e2, e2)), op(e2, e2))) & (e3 = op(op(e2, e2), op(e2, e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax88)).
fof(f1823, plain, (~ spl10_241 | ~ spl10_235 | ~ spl10_223 | ~ spl10_123), inference(avatar_split_clause, [], [f603, f1157, f1700, f1771, f1819])).
fof(f603, plain, (~ (op(e0, e0) = e2) | ~ (e4 = op(e0, op(e0, e0))) | ~ (e1 = op(op(e0, op(e0, e0)), op(e0, e0))) | ~ (e3 = op(op(e0, e0), op(e0, e0)))), inference(cnf_transformation, [], [f206])).
fof(f206, plain, (~ (op(e0, e0) = e2) | ~ (e4 = op(e0, op(e0, e0))) | ~ (e1 = op(op(e0, op(e0, e0)), op(e0, e0))) | ~ (e3 = op(op(e0, e0), op(e0, e0)))), inference(ennf_transformation, [], [f86])).
fof(f86, plain, ~ ((op(e0, e0) = e2) & (e4 = op(e0, op(e0, e0))) & (e1 = op(op(e0, op(e0, e0)), op(e0, e0))) & (e3 = op(op(e0, e0), op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax86)).
fof(f1817, plain, (~ spl10_240 | ~ spl10_219 | ~ spl10_196 | ~ spl10_2), inference(avatar_split_clause, [], [f601, f649, f1565, f1681, f1813])).
fof(f601, plain, (~ (e1 = op(e4, e4)) | ~ (e2 = op(e4, op(e4, e4))) | ~ (e0 = op(op(e4, op(e4, e4)), op(e4, e4))) | ~ (e3 = op(op(e4, e4), op(e4, e4)))), inference(cnf_transformation, [], [f204])).
fof(f204, plain, (~ (e1 = op(e4, e4)) | ~ (e2 = op(e4, op(e4, e4))) | ~ (e0 = op(op(e4, op(e4, e4)), op(e4, e4))) | ~ (e3 = op(op(e4, e4), op(e4, e4)))), inference(ennf_transformation, [], [f84])).
fof(f84, plain, ~ ((e1 = op(e4, e4)) & (e2 = op(e4, op(e4, e4))) & (e0 = op(op(e4, op(e4, e4)), op(e4, e4))) & (e3 = op(op(e4, e4), op(e4, e4)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax84)).
fof(f1805, plain, (~ spl10_238 | ~ spl10_231 | ~ spl10_201 | ~ spl10_93), inference(avatar_split_clause, [], [f597, f1031, f1588, f1749, f1801])).
fof(f597, plain, (~ (e2 = op(e1, e1)) | ~ (e4 = op(e1, op(e1, e1))) | ~ (e0 = op(op(e1, op(e1, e1)), op(e1, e1))) | ~ (e3 = op(op(e1, e1), op(e1, e1)))), inference(cnf_transformation, [], [f200])).
fof(f200, plain, (~ (e2 = op(e1, e1)) | ~ (e4 = op(e1, op(e1, e1))) | ~ (e0 = op(op(e1, op(e1, e1)), op(e1, e1))) | ~ (e3 = op(op(e1, e1), op(e1, e1)))), inference(ennf_transformation, [], [f80])).
fof(f80, plain, ~ ((e2 = op(e1, e1)) & (e4 = op(e1, op(e1, e1))) & (e0 = op(op(e1, op(e1, e1)), op(e1, e1))) & (e3 = op(op(e1, e1), op(e1, e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax80)).
fof(f1804, plain, (~ spl10_238 | ~ spl10_231 | ~ spl10_207 | ~ spl10_95), inference(avatar_split_clause, [], [f596, f1039, f1617, f1749, f1801])).
fof(f596, plain, (~ (e4 = op(e1, e1)) | ~ (e2 = op(e1, op(e1, e1))) | ~ (e0 = op(op(e1, op(e1, e1)), op(e1, e1))) | ~ (e3 = op(op(e1, e1), op(e1, e1)))), inference(cnf_transformation, [], [f199])).
fof(f199, plain, (~ (e4 = op(e1, e1)) | ~ (e2 = op(e1, op(e1, e1))) | ~ (e0 = op(op(e1, op(e1, e1)), op(e1, e1))) | ~ (e3 = op(op(e1, e1), op(e1, e1)))), inference(ennf_transformation, [], [f79])).
fof(f79, plain, ~ ((e4 = op(e1, e1)) & (e2 = op(e1, op(e1, e1))) & (e0 = op(op(e1, op(e1, e1)), op(e1, e1))) & (e3 = op(op(e1, e1), op(e1, e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax79)).
fof(f1799, plain, (~ spl10_232 | ~ spl10_213 | ~ spl10_203 | ~ spl10_31), inference(avatar_split_clause, [], [f595, f771, f1597, f1651, f1755])).
fof(f595, plain, (~ (e0 = op(e3, e3)) | ~ (e1 = op(e3, op(e3, e3))) | ~ (e4 = op(op(e3, op(e3, e3)), op(e3, e3))) | ~ (e2 = op(op(e3, e3), op(e3, e3)))), inference(cnf_transformation, [], [f198])).
fof(f198, plain, (~ (e0 = op(e3, e3)) | ~ (e1 = op(e3, op(e3, e3))) | ~ (e4 = op(op(e3, op(e3, e3)), op(e3, e3))) | ~ (e2 = op(op(e3, e3), op(e3, e3)))), inference(ennf_transformation, [], [f78])).
fof(f78, plain, ~ ((e0 = op(e3, e3)) & (e1 = op(e3, op(e3, e3))) & (e4 = op(op(e3, op(e3, e3)), op(e3, e3))) & (e2 = op(op(e3, e3), op(e3, e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax78)).
fof(f1798, plain, (~ spl10_232 | ~ spl10_213 | ~ spl10_224 | ~ spl10_32), inference(avatar_split_clause, [], [f594, f775, f1705, f1651, f1755])).
fof(f594, plain, (~ (e1 = op(e3, e3)) | ~ (e0 = op(e3, op(e3, e3))) | ~ (e4 = op(op(e3, op(e3, e3)), op(e3, e3))) | ~ (e2 = op(op(e3, e3), op(e3, e3)))), inference(cnf_transformation, [], [f197])).
fof(f197, plain, (~ (e1 = op(e3, e3)) | ~ (e0 = op(e3, op(e3, e3))) | ~ (e4 = op(op(e3, op(e3, e3)), op(e3, e3))) | ~ (e2 = op(op(e3, e3), op(e3, e3)))), inference(ennf_transformation, [], [f77])).
fof(f77, plain, ~ ((e1 = op(e3, e3)) & (e0 = op(e3, op(e3, e3))) & (e4 = op(op(e3, op(e3, e3)), op(e3, e3))) & (e2 = op(op(e3, e3), op(e3, e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax77)).
fof(f1797, plain, (~ spl10_230 | ~ spl10_211 | ~ spl10_200 | ~ spl10_91), inference(avatar_split_clause, [], [f593, f1023, f1583, f1639, f1745])).
fof(f593, plain, (~ (e0 = op(e1, e1)) | ~ (e3 = op(e1, op(e1, e1))) | ~ (e4 = op(op(e1, op(e1, e1)), op(e1, e1))) | ~ (e2 = op(op(e1, e1), op(e1, e1)))), inference(cnf_transformation, [], [f196])).
fof(f196, plain, (~ (e0 = op(e1, e1)) | ~ (e3 = op(e1, op(e1, e1))) | ~ (e4 = op(op(e1, op(e1, e1)), op(e1, e1))) | ~ (e2 = op(op(e1, e1), op(e1, e1)))), inference(ennf_transformation, [], [f76])).
fof(f76, plain, ~ ((e0 = op(e1, e1)) & (e3 = op(e1, op(e1, e1))) & (e4 = op(op(e1, op(e1, e1)), op(e1, e1))) & (e2 = op(op(e1, e1), op(e1, e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax76)).
fof(f1796, plain, (~ spl10_230 | ~ spl10_211 | ~ spl10_237 | ~ spl10_94), inference(avatar_split_clause, [], [f592, f1035, f1787, f1639, f1745])).
fof(f592, plain, (~ (e3 = op(e1, e1)) | ~ (e0 = op(e1, op(e1, e1))) | ~ (e4 = op(op(e1, op(e1, e1)), op(e1, e1))) | ~ (e2 = op(op(e1, e1), op(e1, e1)))), inference(cnf_transformation, [], [f195])).
fof(f195, plain, (~ (e3 = op(e1, e1)) | ~ (e0 = op(e1, op(e1, e1))) | ~ (e4 = op(op(e1, op(e1, e1)), op(e1, e1))) | ~ (e2 = op(op(e1, e1), op(e1, e1)))), inference(ennf_transformation, [], [f75])).
fof(f75, plain, ~ ((e3 = op(e1, e1)) & (e0 = op(e1, op(e1, e1))) & (e4 = op(op(e1, op(e1, e1)), op(e1, e1))) & (e2 = op(op(e1, e1), op(e1, e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax75)).
fof(f1794, plain, (~ spl10_234 | ~ spl10_229 | ~ spl10_236 | ~ spl10_124), inference(avatar_split_clause, [], [f590, f1161, f1781, f1735, f1767])).
fof(f590, plain, (~ (op(e0, e0) = e3) | ~ (e1 = op(e0, op(e0, e0))) | ~ (e4 = op(op(e0, op(e0, e0)), op(e0, e0))) | ~ (e2 = op(op(e0, e0), op(e0, e0)))), inference(cnf_transformation, [], [f193])).
fof(f193, plain, (~ (op(e0, e0) = e3) | ~ (e1 = op(e0, op(e0, e0))) | ~ (e4 = op(op(e0, op(e0, e0)), op(e0, e0))) | ~ (e2 = op(op(e0, e0), op(e0, e0)))), inference(ennf_transformation, [], [f73])).
fof(f73, plain, ~ ((op(e0, e0) = e3) & (e1 = op(e0, op(e0, e0))) & (e4 = op(op(e0, op(e0, e0)), op(e0, e0))) & (e2 = op(op(e0, e0), op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax73)).
fof(f1793, plain, (~ spl10_233 | ~ spl10_210 | ~ spl10_205 | ~ spl10_1), inference(avatar_split_clause, [], [f589, f645, f1607, f1633, f1761])).
fof(f589, plain, (~ (e0 = op(e4, e4)) | ~ (e1 = op(e4, op(e4, e4))) | ~ (e3 = op(op(e4, op(e4, e4)), op(e4, e4))) | ~ (e2 = op(op(e4, e4), op(e4, e4)))), inference(cnf_transformation, [], [f192])).
fof(f192, plain, (~ (e0 = op(e4, e4)) | ~ (e1 = op(e4, op(e4, e4))) | ~ (e3 = op(op(e4, op(e4, e4)), op(e4, e4))) | ~ (e2 = op(op(e4, e4), op(e4, e4)))), inference(ennf_transformation, [], [f72])).
fof(f72, plain, ~ ((e0 = op(e4, e4)) & (e1 = op(e4, op(e4, e4))) & (e3 = op(op(e4, op(e4, e4)), op(e4, e4))) & (e2 = op(op(e4, e4), op(e4, e4)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax72)).
fof(f1792, plain, (~ spl10_233 | ~ spl10_210 | ~ spl10_225 | ~ spl10_2), inference(avatar_split_clause, [], [f588, f649, f1711, f1633, f1761])).
fof(f588, plain, (~ (e1 = op(e4, e4)) | ~ (e0 = op(e4, op(e4, e4))) | ~ (e3 = op(op(e4, op(e4, e4)), op(e4, e4))) | ~ (e2 = op(op(e4, e4), op(e4, e4)))), inference(cnf_transformation, [], [f191])).
fof(f191, plain, (~ (e1 = op(e4, e4)) | ~ (e0 = op(e4, op(e4, e4))) | ~ (e3 = op(op(e4, op(e4, e4)), op(e4, e4))) | ~ (e2 = op(op(e4, e4), op(e4, e4)))), inference(ennf_transformation, [], [f71])).
fof(f71, plain, ~ ((e1 = op(e4, e4)) & (e0 = op(e4, op(e4, e4))) & (e3 = op(op(e4, op(e4, e4)), op(e4, e4))) & (e2 = op(op(e4, e4), op(e4, e4)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax71)).
fof(f1791, plain, (~ spl10_230 | ~ spl10_206 | ~ spl10_201 | ~ spl10_91), inference(avatar_split_clause, [], [f587, f1023, f1588, f1613, f1745])).
fof(f587, plain, (~ (e0 = op(e1, e1)) | ~ (e4 = op(e1, op(e1, e1))) | ~ (e3 = op(op(e1, op(e1, e1)), op(e1, e1))) | ~ (e2 = op(op(e1, e1), op(e1, e1)))), inference(cnf_transformation, [], [f190])).
fof(f190, plain, (~ (e0 = op(e1, e1)) | ~ (e4 = op(e1, op(e1, e1))) | ~ (e3 = op(op(e1, op(e1, e1)), op(e1, e1))) | ~ (e2 = op(op(e1, e1), op(e1, e1)))), inference(ennf_transformation, [], [f70])).
fof(f70, plain, ~ ((e0 = op(e1, e1)) & (e4 = op(e1, op(e1, e1))) & (e3 = op(op(e1, op(e1, e1)), op(e1, e1))) & (e2 = op(op(e1, e1), op(e1, e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax70)).
fof(f1785, plain, (~ spl10_234 | ~ spl10_226 | ~ spl10_223 | ~ spl10_122), inference(avatar_split_clause, [], [f585, f1153, f1700, f1717, f1767])).
fof(f585, plain, (~ (op(e0, e0) = e1) | ~ (e4 = op(e0, op(e0, e0))) | ~ (e3 = op(op(e0, op(e0, e0)), op(e0, e0))) | ~ (e2 = op(op(e0, e0), op(e0, e0)))), inference(cnf_transformation, [], [f188])).
fof(f188, plain, (~ (op(e0, e0) = e1) | ~ (e4 = op(e0, op(e0, e0))) | ~ (e3 = op(op(e0, op(e0, e0)), op(e0, e0))) | ~ (e2 = op(op(e0, e0), op(e0, e0)))), inference(ennf_transformation, [], [f68])).
fof(f68, plain, ~ ((op(e0, e0) = e1) & (e4 = op(e0, op(e0, e0))) & (e3 = op(op(e0, op(e0, e0)), op(e0, e0))) & (e2 = op(op(e0, e0), op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax68)).
fof(f1784, plain, (~ spl10_234 | ~ spl10_226 | ~ spl10_236 | ~ spl10_125), inference(avatar_split_clause, [], [f584, f1165, f1781, f1717, f1767])).
fof(f584, plain, (~ (op(e0, e0) = e4) | ~ (e1 = op(e0, op(e0, e0))) | ~ (e3 = op(op(e0, op(e0, e0)), op(e0, e0))) | ~ (e2 = op(op(e0, e0), op(e0, e0)))), inference(cnf_transformation, [], [f187])).
fof(f187, plain, (~ (op(e0, e0) = e4) | ~ (e1 = op(e0, op(e0, e0))) | ~ (e3 = op(op(e0, op(e0, e0)), op(e0, e0))) | ~ (e2 = op(op(e0, e0), op(e0, e0)))), inference(ennf_transformation, [], [f67])).
fof(f67, plain, ~ ((op(e0, e0) = e4) & (e1 = op(e0, op(e0, e0))) & (e3 = op(op(e0, op(e0, e0)), op(e0, e0))) & (e2 = op(op(e0, e0), op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax67)).
fof(f1778, plain, (~ spl10_233 | ~ spl10_195 | ~ spl10_225 | ~ spl10_4), inference(avatar_split_clause, [], [f582, f657, f1711, f1561, f1761])).
fof(f582, plain, (~ (e3 = op(e4, e4)) | ~ (e0 = op(e4, op(e4, e4))) | ~ (e1 = op(op(e4, op(e4, e4)), op(e4, e4))) | ~ (e2 = op(op(e4, e4), op(e4, e4)))), inference(cnf_transformation, [], [f185])).
fof(f185, plain, (~ (e3 = op(e4, e4)) | ~ (e0 = op(e4, op(e4, e4))) | ~ (e1 = op(op(e4, op(e4, e4)), op(e4, e4))) | ~ (e2 = op(op(e4, e4), op(e4, e4)))), inference(ennf_transformation, [], [f65])).
fof(f65, plain, ~ ((e3 = op(e4, e4)) & (e0 = op(e4, op(e4, e4))) & (e1 = op(op(e4, op(e4, e4)), op(e4, e4))) & (e2 = op(op(e4, e4), op(e4, e4)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax65)).
fof(f1777, plain, (~ spl10_232 | ~ spl10_191 | ~ spl10_193 | ~ spl10_31), inference(avatar_split_clause, [], [f581, f771, f1552, f1543, f1755])).
fof(f581, plain, (~ (e0 = op(e3, e3)) | ~ (e4 = op(e3, op(e3, e3))) | ~ (e1 = op(op(e3, op(e3, e3)), op(e3, e3))) | ~ (e2 = op(op(e3, e3), op(e3, e3)))), inference(cnf_transformation, [], [f184])).
fof(f184, plain, (~ (e0 = op(e3, e3)) | ~ (e4 = op(e3, op(e3, e3))) | ~ (e1 = op(op(e3, op(e3, e3)), op(e3, e3))) | ~ (e2 = op(op(e3, e3), op(e3, e3)))), inference(ennf_transformation, [], [f64])).
fof(f64, plain, ~ ((e0 = op(e3, e3)) & (e4 = op(e3, op(e3, e3))) & (e1 = op(op(e3, op(e3, e3)), op(e3, e3))) & (e2 = op(op(e3, e3), op(e3, e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax64)).
fof(f1775, plain, (~ spl10_234 | ~ spl10_235 | ~ spl10_223 | ~ spl10_124), inference(avatar_split_clause, [], [f579, f1161, f1700, f1771, f1767])).
fof(f579, plain, (~ (op(e0, e0) = e3) | ~ (e4 = op(e0, op(e0, e0))) | ~ (e1 = op(op(e0, op(e0, e0)), op(e0, e0))) | ~ (e2 = op(op(e0, e0), op(e0, e0)))), inference(cnf_transformation, [], [f182])).
fof(f182, plain, (~ (op(e0, e0) = e3) | ~ (e4 = op(e0, op(e0, e0))) | ~ (e1 = op(op(e0, op(e0, e0)), op(e0, e0))) | ~ (e2 = op(op(e0, e0), op(e0, e0)))), inference(ennf_transformation, [], [f62])).
fof(f62, plain, ~ ((op(e0, e0) = e3) & (e4 = op(e0, op(e0, e0))) & (e1 = op(op(e0, op(e0, e0)), op(e0, e0))) & (e2 = op(op(e0, e0), op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax62)).
fof(f1764, plain, (~ spl10_233 | ~ spl10_219 | ~ spl10_205 | ~ spl10_4), inference(avatar_split_clause, [], [f576, f657, f1607, f1681, f1761])).
fof(f576, plain, (~ (e3 = op(e4, e4)) | ~ (e1 = op(e4, op(e4, e4))) | ~ (e0 = op(op(e4, op(e4, e4)), op(e4, e4))) | ~ (e2 = op(op(e4, e4), op(e4, e4)))), inference(cnf_transformation, [], [f179])).
fof(f179, plain, (~ (e3 = op(e4, e4)) | ~ (e1 = op(e4, op(e4, e4))) | ~ (e0 = op(op(e4, op(e4, e4)), op(e4, e4))) | ~ (e2 = op(op(e4, e4), op(e4, e4)))), inference(ennf_transformation, [], [f59])).
fof(f59, plain, ~ ((e3 = op(e4, e4)) & (e1 = op(e4, op(e4, e4))) & (e0 = op(op(e4, op(e4, e4)), op(e4, e4))) & (e2 = op(op(e4, e4), op(e4, e4)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax59)).
fof(f1759, plain, (~ spl10_232 | ~ spl10_217 | ~ spl10_193 | ~ spl10_32), inference(avatar_split_clause, [], [f575, f775, f1552, f1671, f1755])).
fof(f575, plain, (~ (e1 = op(e3, e3)) | ~ (e4 = op(e3, op(e3, e3))) | ~ (e0 = op(op(e3, op(e3, e3)), op(e3, e3))) | ~ (e2 = op(op(e3, e3), op(e3, e3)))), inference(cnf_transformation, [], [f178])).
fof(f178, plain, (~ (e1 = op(e3, e3)) | ~ (e4 = op(e3, op(e3, e3))) | ~ (e0 = op(op(e3, op(e3, e3)), op(e3, e3))) | ~ (e2 = op(op(e3, e3), op(e3, e3)))), inference(ennf_transformation, [], [f58])).
fof(f58, plain, ~ ((e1 = op(e3, e3)) & (e4 = op(e3, op(e3, e3))) & (e0 = op(op(e3, op(e3, e3)), op(e3, e3))) & (e2 = op(op(e3, e3), op(e3, e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax58)).
fof(f1758, plain, (~ spl10_232 | ~ spl10_217 | ~ spl10_203 | ~ spl10_35), inference(avatar_split_clause, [], [f574, f787, f1597, f1671, f1755])).
fof(f574, plain, (~ (e4 = op(e3, e3)) | ~ (e1 = op(e3, op(e3, e3))) | ~ (e0 = op(op(e3, op(e3, e3)), op(e3, e3))) | ~ (e2 = op(op(e3, e3), op(e3, e3)))), inference(cnf_transformation, [], [f177])).
fof(f177, plain, (~ (e4 = op(e3, e3)) | ~ (e1 = op(e3, op(e3, e3))) | ~ (e0 = op(op(e3, op(e3, e3)), op(e3, e3))) | ~ (e2 = op(op(e3, e3), op(e3, e3)))), inference(ennf_transformation, [], [f57])).
fof(f57, plain, ~ ((e4 = op(e3, e3)) & (e1 = op(e3, op(e3, e3))) & (e0 = op(op(e3, op(e3, e3)), op(e3, e3))) & (e2 = op(op(e3, e3), op(e3, e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax57)).
fof(f1752, plain, (~ spl10_230 | ~ spl10_231 | ~ spl10_200 | ~ spl10_95), inference(avatar_split_clause, [], [f572, f1039, f1583, f1749, f1745])).
fof(f572, plain, (~ (e4 = op(e1, e1)) | ~ (e3 = op(e1, op(e1, e1))) | ~ (e0 = op(op(e1, op(e1, e1)), op(e1, e1))) | ~ (e2 = op(op(e1, e1), op(e1, e1)))), inference(cnf_transformation, [], [f175])).
fof(f175, plain, (~ (e4 = op(e1, e1)) | ~ (e3 = op(e1, op(e1, e1))) | ~ (e0 = op(op(e1, op(e1, e1)), op(e1, e1))) | ~ (e2 = op(op(e1, e1), op(e1, e1)))), inference(ennf_transformation, [], [f55])).
fof(f55, plain, ~ ((e4 = op(e1, e1)) & (e3 = op(e1, op(e1, e1))) & (e0 = op(op(e1, op(e1, e1)), op(e1, e1))) & (e2 = op(op(e1, e1), op(e1, e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax55)).
fof(f1743, plain, (~ spl10_216 | ~ spl10_213 | ~ spl10_192 | ~ spl10_31), inference(avatar_split_clause, [], [f571, f771, f1547, f1651, f1667])).
fof(f571, plain, (~ (e0 = op(e3, e3)) | ~ (e2 = op(e3, op(e3, e3))) | ~ (e4 = op(op(e3, op(e3, e3)), op(e3, e3))) | ~ (e1 = op(op(e3, e3), op(e3, e3)))), inference(cnf_transformation, [], [f174])).
fof(f174, plain, (~ (e0 = op(e3, e3)) | ~ (e2 = op(e3, op(e3, e3))) | ~ (e4 = op(op(e3, op(e3, e3)), op(e3, e3))) | ~ (e1 = op(op(e3, e3), op(e3, e3)))), inference(ennf_transformation, [], [f54])).
fof(f54, plain, ~ ((e0 = op(e3, e3)) & (e2 = op(e3, op(e3, e3))) & (e4 = op(op(e3, op(e3, e3)), op(e3, e3))) & (e1 = op(op(e3, e3), op(e3, e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax54)).
fof(f1742, plain, (~ spl10_216 | ~ spl10_213 | ~ spl10_224 | ~ spl10_33), inference(avatar_split_clause, [], [f570, f779, f1705, f1651, f1667])).
fof(f570, plain, (~ (e2 = op(e3, e3)) | ~ (e0 = op(e3, op(e3, e3))) | ~ (e4 = op(op(e3, op(e3, e3)), op(e3, e3))) | ~ (e1 = op(op(e3, e3), op(e3, e3)))), inference(cnf_transformation, [], [f173])).
fof(f173, plain, (~ (e2 = op(e3, e3)) | ~ (e0 = op(e3, op(e3, e3))) | ~ (e4 = op(op(e3, op(e3, e3)), op(e3, e3))) | ~ (e1 = op(op(e3, e3), op(e3, e3)))), inference(ennf_transformation, [], [f53])).
fof(f53, plain, ~ ((e2 = op(e3, e3)) & (e0 = op(e3, op(e3, e3))) & (e4 = op(op(e3, op(e3, e3)), op(e3, e3))) & (e1 = op(op(e3, e3), op(e3, e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax53)).
fof(f1741, plain, (~ spl10_214 | ~ spl10_212 | ~ spl10_188 | ~ spl10_61), inference(avatar_split_clause, [], [f569, f897, f1529, f1645, f1657])).
fof(f569, plain, (~ (e0 = op(e2, e2)) | ~ (e3 = op(e2, op(e2, e2))) | ~ (e4 = op(op(e2, op(e2, e2)), op(e2, e2))) | ~ (e1 = op(op(e2, e2), op(e2, e2)))), inference(cnf_transformation, [], [f172])).
fof(f172, plain, (~ (e0 = op(e2, e2)) | ~ (e3 = op(e2, op(e2, e2))) | ~ (e4 = op(op(e2, op(e2, e2)), op(e2, e2))) | ~ (e1 = op(op(e2, e2), op(e2, e2)))), inference(ennf_transformation, [], [f52])).
fof(f52, plain, ~ ((e0 = op(e2, e2)) & (e3 = op(e2, op(e2, e2))) & (e4 = op(op(e2, op(e2, e2)), op(e2, e2))) & (e1 = op(op(e2, e2), op(e2, e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax52)).
fof(f1732, plain, (~ spl10_218 | ~ spl10_210 | ~ spl10_225 | ~ spl10_3), inference(avatar_split_clause, [], [f564, f653, f1711, f1633, f1677])).
fof(f564, plain, (~ (e2 = op(e4, e4)) | ~ (e0 = op(e4, op(e4, e4))) | ~ (e3 = op(op(e4, op(e4, e4)), op(e4, e4))) | ~ (e1 = op(op(e4, e4), op(e4, e4)))), inference(cnf_transformation, [], [f167])).
fof(f167, plain, (~ (e2 = op(e4, e4)) | ~ (e0 = op(e4, op(e4, e4))) | ~ (e3 = op(op(e4, op(e4, e4)), op(e4, e4))) | ~ (e1 = op(op(e4, e4), op(e4, e4)))), inference(ennf_transformation, [], [f47])).
fof(f47, plain, ~ ((e2 = op(e4, e4)) & (e0 = op(e4, op(e4, e4))) & (e3 = op(op(e4, op(e4, e4)), op(e4, e4))) & (e1 = op(op(e4, e4), op(e4, e4)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax47)).
fof(f1730, plain, (~ spl10_214 | ~ spl10_208 | ~ spl10_228 | ~ spl10_65), inference(avatar_split_clause, [], [f562, f913, f1727, f1623, f1657])).
fof(f562, plain, (~ (e4 = op(e2, e2)) | ~ (e0 = op(e2, op(e2, e2))) | ~ (e3 = op(op(e2, op(e2, e2)), op(e2, e2))) | ~ (e1 = op(op(e2, e2), op(e2, e2)))), inference(cnf_transformation, [], [f165])).
fof(f165, plain, (~ (e4 = op(e2, e2)) | ~ (e0 = op(e2, op(e2, e2))) | ~ (e3 = op(op(e2, op(e2, e2)), op(e2, e2))) | ~ (e1 = op(op(e2, e2), op(e2, e2)))), inference(ennf_transformation, [], [f45])).
fof(f45, plain, ~ ((e4 = op(e2, e2)) & (e0 = op(e2, op(e2, e2))) & (e3 = op(op(e2, op(e2, e2)), op(e2, e2))) & (e1 = op(op(e2, e2), op(e2, e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax45)).
fof(f1715, plain, (~ spl10_218 | ~ spl10_204 | ~ spl10_197 | ~ spl10_1), inference(avatar_split_clause, [], [f559, f645, f1570, f1603, f1677])).
fof(f559, plain, (~ (e0 = op(e4, e4)) | ~ (e3 = op(e4, op(e4, e4))) | ~ (e2 = op(op(e4, op(e4, e4)), op(e4, e4))) | ~ (e1 = op(op(e4, e4), op(e4, e4)))), inference(cnf_transformation, [], [f162])).
fof(f162, plain, (~ (e0 = op(e4, e4)) | ~ (e3 = op(e4, op(e4, e4))) | ~ (e2 = op(op(e4, op(e4, e4)), op(e4, e4))) | ~ (e1 = op(op(e4, e4), op(e4, e4)))), inference(ennf_transformation, [], [f42])).
fof(f42, plain, ~ ((e0 = op(e4, e4)) & (e3 = op(e4, op(e4, e4))) & (e2 = op(op(e4, op(e4, e4)), op(e4, e4))) & (e1 = op(op(e4, e4), op(e4, e4)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax42)).
fof(f1714, plain, (~ spl10_218 | ~ spl10_204 | ~ spl10_225 | ~ spl10_4), inference(avatar_split_clause, [], [f558, f657, f1711, f1603, f1677])).
fof(f558, plain, (~ (e3 = op(e4, e4)) | ~ (e0 = op(e4, op(e4, e4))) | ~ (e2 = op(op(e4, op(e4, e4)), op(e4, e4))) | ~ (e1 = op(op(e4, e4), op(e4, e4)))), inference(cnf_transformation, [], [f161])).
fof(f161, plain, (~ (e3 = op(e4, e4)) | ~ (e0 = op(e4, op(e4, e4))) | ~ (e2 = op(op(e4, op(e4, e4)), op(e4, e4))) | ~ (e1 = op(op(e4, e4), op(e4, e4)))), inference(ennf_transformation, [], [f41])).
fof(f41, plain, ~ ((e3 = op(e4, e4)) & (e0 = op(e4, op(e4, e4))) & (e2 = op(op(e4, op(e4, e4)), op(e4, e4))) & (e1 = op(op(e4, e4), op(e4, e4)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax41)).
fof(f1708, plain, (~ spl10_216 | ~ spl10_202 | ~ spl10_224 | ~ spl10_35), inference(avatar_split_clause, [], [f556, f787, f1705, f1593, f1667])).
fof(f556, plain, (~ (e4 = op(e3, e3)) | ~ (e0 = op(e3, op(e3, e3))) | ~ (e2 = op(op(e3, op(e3, e3)), op(e3, e3))) | ~ (e1 = op(op(e3, e3), op(e3, e3)))), inference(cnf_transformation, [], [f159])).
fof(f159, plain, (~ (e4 = op(e3, e3)) | ~ (e0 = op(e3, op(e3, e3))) | ~ (e2 = op(op(e3, op(e3, e3)), op(e3, e3))) | ~ (e1 = op(op(e3, e3), op(e3, e3)))), inference(ennf_transformation, [], [f39])).
fof(f39, plain, ~ ((e4 = op(e3, e3)) & (e0 = op(e3, op(e3, e3))) & (e2 = op(op(e3, op(e3, e3)), op(e3, e3))) & (e1 = op(op(e3, e3), op(e3, e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax39)).
fof(f1703, plain, (~ spl10_220 | ~ spl10_221 | ~ spl10_223 | ~ spl10_124), inference(avatar_split_clause, [], [f555, f1161, f1700, f1691, f1687])).
fof(f555, plain, (~ (op(e0, e0) = e3) | ~ (e4 = op(e0, op(e0, e0))) | ~ (e2 = op(op(e0, op(e0, e0)), op(e0, e0))) | ~ (e1 = op(op(e0, e0), op(e0, e0)))), inference(cnf_transformation, [], [f158])).
fof(f158, plain, (~ (op(e0, e0) = e3) | ~ (e4 = op(e0, op(e0, e0))) | ~ (e2 = op(op(e0, op(e0, e0)), op(e0, e0))) | ~ (e1 = op(op(e0, e0), op(e0, e0)))), inference(ennf_transformation, [], [f38])).
fof(f38, plain, ~ ((op(e0, e0) = e3) & (e4 = op(e0, op(e0, e0))) & (e2 = op(op(e0, op(e0, e0)), op(e0, e0))) & (e1 = op(op(e0, e0), op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax38)).
fof(f1684, plain, (~ spl10_218 | ~ spl10_219 | ~ spl10_196 | ~ spl10_4), inference(avatar_split_clause, [], [f552, f657, f1565, f1681, f1677])).
fof(f552, plain, (~ (e3 = op(e4, e4)) | ~ (e2 = op(e4, op(e4, e4))) | ~ (e0 = op(op(e4, op(e4, e4)), op(e4, e4))) | ~ (e1 = op(op(e4, e4), op(e4, e4)))), inference(cnf_transformation, [], [f155])).
fof(f155, plain, (~ (e3 = op(e4, e4)) | ~ (e2 = op(e4, op(e4, e4))) | ~ (e0 = op(op(e4, op(e4, e4)), op(e4, e4))) | ~ (e1 = op(op(e4, e4), op(e4, e4)))), inference(ennf_transformation, [], [f35])).
fof(f35, plain, ~ ((e3 = op(e4, e4)) & (e2 = op(e4, op(e4, e4))) & (e0 = op(op(e4, op(e4, e4)), op(e4, e4))) & (e1 = op(op(e4, e4), op(e4, e4)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax35)).
fof(f1675, plain, (~ spl10_216 | ~ spl10_217 | ~ spl10_193 | ~ spl10_33), inference(avatar_split_clause, [], [f551, f779, f1552, f1671, f1667])).
fof(f551, plain, (~ (e2 = op(e3, e3)) | ~ (e4 = op(e3, op(e3, e3))) | ~ (e0 = op(op(e3, op(e3, e3)), op(e3, e3))) | ~ (e1 = op(op(e3, e3), op(e3, e3)))), inference(cnf_transformation, [], [f154])).
fof(f154, plain, (~ (e2 = op(e3, e3)) | ~ (e4 = op(e3, op(e3, e3))) | ~ (e0 = op(op(e3, op(e3, e3)), op(e3, e3))) | ~ (e1 = op(op(e3, e3), op(e3, e3)))), inference(ennf_transformation, [], [f34])).
fof(f34, plain, ~ ((e2 = op(e3, e3)) & (e4 = op(e3, op(e3, e3))) & (e0 = op(op(e3, op(e3, e3)), op(e3, e3))) & (e1 = op(op(e3, e3), op(e3, e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax34)).
fof(f1674, plain, (~ spl10_216 | ~ spl10_217 | ~ spl10_192 | ~ spl10_35), inference(avatar_split_clause, [], [f550, f787, f1547, f1671, f1667])).
fof(f550, plain, (~ (e4 = op(e3, e3)) | ~ (e2 = op(e3, op(e3, e3))) | ~ (e0 = op(op(e3, op(e3, e3)), op(e3, e3))) | ~ (e1 = op(op(e3, e3), op(e3, e3)))), inference(cnf_transformation, [], [f153])).
fof(f153, plain, (~ (e4 = op(e3, e3)) | ~ (e2 = op(e3, op(e3, e3))) | ~ (e0 = op(op(e3, op(e3, e3)), op(e3, e3))) | ~ (e1 = op(op(e3, e3), op(e3, e3)))), inference(ennf_transformation, [], [f33])).
fof(f33, plain, ~ ((e4 = op(e3, e3)) & (e2 = op(e3, op(e3, e3))) & (e0 = op(op(e3, op(e3, e3)), op(e3, e3))) & (e1 = op(op(e3, e3), op(e3, e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax33)).
fof(f1665, plain, (~ spl10_214 | ~ spl10_215 | ~ spl10_189 | ~ spl10_64), inference(avatar_split_clause, [], [f549, f909, f1534, f1661, f1657])).
fof(f549, plain, (~ (e3 = op(e2, e2)) | ~ (e4 = op(e2, op(e2, e2))) | ~ (e0 = op(op(e2, op(e2, e2)), op(e2, e2))) | ~ (e1 = op(op(e2, e2), op(e2, e2)))), inference(cnf_transformation, [], [f152])).
fof(f152, plain, (~ (e3 = op(e2, e2)) | ~ (e4 = op(e2, op(e2, e2))) | ~ (e0 = op(op(e2, op(e2, e2)), op(e2, e2))) | ~ (e1 = op(op(e2, e2), op(e2, e2)))), inference(ennf_transformation, [], [f32])).
fof(f32, plain, ~ ((e3 = op(e2, e2)) & (e4 = op(e2, op(e2, e2))) & (e0 = op(op(e2, op(e2, e2)), op(e2, e2))) & (e1 = op(op(e2, e2), op(e2, e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax32)).
fof(f1664, plain, (~ spl10_214 | ~ spl10_215 | ~ spl10_188 | ~ spl10_65), inference(avatar_split_clause, [], [f548, f913, f1529, f1661, f1657])).
fof(f548, plain, (~ (e4 = op(e2, e2)) | ~ (e3 = op(e2, op(e2, e2))) | ~ (e0 = op(op(e2, op(e2, e2)), op(e2, e2))) | ~ (e1 = op(op(e2, e2), op(e2, e2)))), inference(cnf_transformation, [], [f151])).
fof(f151, plain, (~ (e4 = op(e2, e2)) | ~ (e3 = op(e2, op(e2, e2))) | ~ (e0 = op(op(e2, op(e2, e2)), op(e2, e2))) | ~ (e1 = op(op(e2, e2), op(e2, e2)))), inference(ennf_transformation, [], [f31])).
fof(f31, plain, ~ ((e4 = op(e2, e2)) & (e3 = op(e2, op(e2, e2))) & (e0 = op(op(e2, op(e2, e2)), op(e2, e2))) & (e1 = op(op(e2, e2), op(e2, e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax31)).
fof(f1655, plain, (~ spl10_190 | ~ spl10_213 | ~ spl10_192 | ~ spl10_32), inference(avatar_split_clause, [], [f547, f775, f1547, f1651, f1539])).
fof(f547, plain, (~ (e1 = op(e3, e3)) | ~ (e2 = op(e3, op(e3, e3))) | ~ (e4 = op(op(e3, op(e3, e3)), op(e3, e3))) | ~ (e0 = op(op(e3, e3), op(e3, e3)))), inference(cnf_transformation, [], [f150])).
fof(f150, plain, (~ (e1 = op(e3, e3)) | ~ (e2 = op(e3, op(e3, e3))) | ~ (e4 = op(op(e3, op(e3, e3)), op(e3, e3))) | ~ (e0 = op(op(e3, e3), op(e3, e3)))), inference(ennf_transformation, [], [f30])).
fof(f30, plain, ~ ((e1 = op(e3, e3)) & (e2 = op(e3, op(e3, e3))) & (e4 = op(op(e3, op(e3, e3)), op(e3, e3))) & (e0 = op(op(e3, e3), op(e3, e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax30)).
fof(f1654, plain, (~ spl10_190 | ~ spl10_213 | ~ spl10_203 | ~ spl10_33), inference(avatar_split_clause, [], [f546, f779, f1597, f1651, f1539])).
fof(f546, plain, (~ (e2 = op(e3, e3)) | ~ (e1 = op(e3, op(e3, e3))) | ~ (e4 = op(op(e3, op(e3, e3)), op(e3, e3))) | ~ (e0 = op(op(e3, e3), op(e3, e3)))), inference(cnf_transformation, [], [f149])).
fof(f149, plain, (~ (e2 = op(e3, e3)) | ~ (e1 = op(e3, op(e3, e3))) | ~ (e4 = op(op(e3, op(e3, e3)), op(e3, e3))) | ~ (e0 = op(op(e3, e3), op(e3, e3)))), inference(ennf_transformation, [], [f29])).
fof(f29, plain, ~ ((e2 = op(e3, e3)) & (e1 = op(e3, op(e3, e3))) & (e4 = op(op(e3, op(e3, e3)), op(e3, e3))) & (e0 = op(op(e3, e3), op(e3, e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax29)).
fof(f1648, plain, (~ spl10_186 | ~ spl10_212 | ~ spl10_209 | ~ spl10_64), inference(avatar_split_clause, [], [f544, f909, f1627, f1645, f1521])).
fof(f544, plain, (~ (e3 = op(e2, e2)) | ~ (e1 = op(e2, op(e2, e2))) | ~ (e4 = op(op(e2, op(e2, e2)), op(e2, e2))) | ~ (e0 = op(op(e2, e2), op(e2, e2)))), inference(cnf_transformation, [], [f147])).
fof(f147, plain, (~ (e3 = op(e2, e2)) | ~ (e1 = op(e2, op(e2, e2))) | ~ (e4 = op(op(e2, op(e2, e2)), op(e2, e2))) | ~ (e0 = op(op(e2, e2), op(e2, e2)))), inference(ennf_transformation, [], [f27])).
fof(f27, plain, ~ ((e3 = op(e2, e2)) & (e1 = op(e2, op(e2, e2))) & (e4 = op(op(e2, op(e2, e2)), op(e2, e2))) & (e0 = op(op(e2, e2), op(e2, e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax27)).
fof(f1643, plain, (~ spl10_198 | ~ spl10_211 | ~ spl10_200 | ~ spl10_93), inference(avatar_split_clause, [], [f543, f1031, f1583, f1639, f1575])).
fof(f543, plain, (~ (e2 = op(e1, e1)) | ~ (e3 = op(e1, op(e1, e1))) | ~ (e4 = op(op(e1, op(e1, e1)), op(e1, e1))) | ~ (e0 = op(op(e1, e1), op(e1, e1)))), inference(cnf_transformation, [], [f146])).
fof(f146, plain, (~ (e2 = op(e1, e1)) | ~ (e3 = op(e1, op(e1, e1))) | ~ (e4 = op(op(e1, op(e1, e1)), op(e1, e1))) | ~ (e0 = op(op(e1, e1), op(e1, e1)))), inference(ennf_transformation, [], [f26])).
fof(f26, plain, ~ ((e2 = op(e1, e1)) & (e3 = op(e1, op(e1, e1))) & (e4 = op(op(e1, op(e1, e1)), op(e1, e1))) & (e0 = op(op(e1, e1), op(e1, e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax26)).
fof(f1642, plain, (~ spl10_198 | ~ spl10_211 | ~ spl10_207 | ~ spl10_94), inference(avatar_split_clause, [], [f542, f1035, f1617, f1639, f1575])).
fof(f542, plain, (~ (e3 = op(e1, e1)) | ~ (e2 = op(e1, op(e1, e1))) | ~ (e4 = op(op(e1, op(e1, e1)), op(e1, e1))) | ~ (e0 = op(op(e1, e1), op(e1, e1)))), inference(cnf_transformation, [], [f145])).
fof(f145, plain, (~ (e3 = op(e1, e1)) | ~ (e2 = op(e1, op(e1, e1))) | ~ (e4 = op(op(e1, op(e1, e1)), op(e1, e1))) | ~ (e0 = op(op(e1, e1), op(e1, e1)))), inference(ennf_transformation, [], [f25])).
fof(f25, plain, ~ ((e3 = op(e1, e1)) & (e2 = op(e1, op(e1, e1))) & (e4 = op(op(e1, op(e1, e1)), op(e1, e1))) & (e0 = op(op(e1, e1), op(e1, e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax25)).
fof(f1636, plain, (~ spl10_194 | ~ spl10_210 | ~ spl10_205 | ~ spl10_3), inference(avatar_split_clause, [], [f540, f653, f1607, f1633, f1557])).
fof(f540, plain, (~ (e2 = op(e4, e4)) | ~ (e1 = op(e4, op(e4, e4))) | ~ (e3 = op(op(e4, op(e4, e4)), op(e4, e4))) | ~ (e0 = op(op(e4, e4), op(e4, e4)))), inference(cnf_transformation, [], [f143])).
fof(f143, plain, (~ (e2 = op(e4, e4)) | ~ (e1 = op(e4, op(e4, e4))) | ~ (e3 = op(op(e4, op(e4, e4)), op(e4, e4))) | ~ (e0 = op(op(e4, e4), op(e4, e4)))), inference(ennf_transformation, [], [f23])).
fof(f23, plain, ~ ((e2 = op(e4, e4)) & (e1 = op(e4, op(e4, e4))) & (e3 = op(op(e4, op(e4, e4)), op(e4, e4))) & (e0 = op(op(e4, e4), op(e4, e4)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax23)).
fof(f1601, plain, (~ spl10_190 | ~ spl10_202 | ~ spl10_193 | ~ spl10_32), inference(avatar_split_clause, [], [f533, f775, f1552, f1593, f1539])).
fof(f533, plain, (~ (e1 = op(e3, e3)) | ~ (e4 = op(e3, op(e3, e3))) | ~ (e2 = op(op(e3, op(e3, e3)), op(e3, e3))) | ~ (e0 = op(op(e3, e3), op(e3, e3)))), inference(cnf_transformation, [], [f136])).
fof(f136, plain, (~ (e1 = op(e3, e3)) | ~ (e4 = op(e3, op(e3, e3))) | ~ (e2 = op(op(e3, op(e3, e3)), op(e3, e3))) | ~ (e0 = op(op(e3, e3), op(e3, e3)))), inference(ennf_transformation, [], [f16])).
fof(f16, plain, ~ ((e1 = op(e3, e3)) & (e4 = op(e3, op(e3, e3))) & (e2 = op(op(e3, op(e3, e3)), op(e3, e3))) & (e0 = op(op(e3, e3), op(e3, e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax16)).
fof(f1568, plain, (~ spl10_194 | ~ spl10_195 | ~ spl10_196 | ~ spl10_4), inference(avatar_split_clause, [], [f528, f657, f1565, f1561, f1557])).
fof(f528, plain, (~ (e3 = op(e4, e4)) | ~ (e2 = op(e4, op(e4, e4))) | ~ (e1 = op(op(e4, op(e4, e4)), op(e4, e4))) | ~ (e0 = op(op(e4, e4), op(e4, e4)))), inference(cnf_transformation, [], [f131])).
fof(f131, plain, (~ (e3 = op(e4, e4)) | ~ (e2 = op(e4, op(e4, e4))) | ~ (e1 = op(op(e4, op(e4, e4)), op(e4, e4))) | ~ (e0 = op(op(e4, e4), op(e4, e4)))), inference(ennf_transformation, [], [f11])).
fof(f11, plain, ~ ((e3 = op(e4, e4)) & (e2 = op(e4, op(e4, e4))) & (e1 = op(op(e4, op(e4, e4)), op(e4, e4))) & (e0 = op(op(e4, e4), op(e4, e4)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax11)).
fof(f1555, plain, (~ spl10_190 | ~ spl10_191 | ~ spl10_193 | ~ spl10_33), inference(avatar_split_clause, [], [f527, f779, f1552, f1543, f1539])).
fof(f527, plain, (~ (e2 = op(e3, e3)) | ~ (e4 = op(e3, op(e3, e3))) | ~ (e1 = op(op(e3, op(e3, e3)), op(e3, e3))) | ~ (e0 = op(op(e3, e3), op(e3, e3)))), inference(cnf_transformation, [], [f130])).
fof(f130, plain, (~ (e2 = op(e3, e3)) | ~ (e4 = op(e3, op(e3, e3))) | ~ (e1 = op(op(e3, op(e3, e3)), op(e3, e3))) | ~ (e0 = op(op(e3, e3), op(e3, e3)))), inference(ennf_transformation, [], [f10])).
fof(f10, plain, ~ ((e2 = op(e3, e3)) & (e4 = op(e3, op(e3, e3))) & (e1 = op(op(e3, op(e3, e3)), op(e3, e3))) & (e0 = op(op(e3, e3), op(e3, e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax10)).
fof(f1550, plain, (~ spl10_190 | ~ spl10_191 | ~ spl10_192 | ~ spl10_35), inference(avatar_split_clause, [], [f526, f787, f1547, f1543, f1539])).
fof(f526, plain, (~ (e4 = op(e3, e3)) | ~ (e2 = op(e3, op(e3, e3))) | ~ (e1 = op(op(e3, op(e3, e3)), op(e3, e3))) | ~ (e0 = op(op(e3, e3), op(e3, e3)))), inference(cnf_transformation, [], [f129])).
fof(f129, plain, (~ (e4 = op(e3, e3)) | ~ (e2 = op(e3, op(e3, e3))) | ~ (e1 = op(op(e3, op(e3, e3)), op(e3, e3))) | ~ (e0 = op(op(e3, e3), op(e3, e3)))), inference(ennf_transformation, [], [f9])).
fof(f9, plain, ~ ((e4 = op(e3, e3)) & (e2 = op(e3, op(e3, e3))) & (e1 = op(op(e3, op(e3, e3)), op(e3, e3))) & (e0 = op(op(e3, e3), op(e3, e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax9)).
fof(f1537, plain, (~ spl10_186 | ~ spl10_187 | ~ spl10_189 | ~ spl10_64), inference(avatar_split_clause, [], [f525, f909, f1534, f1525, f1521])).
fof(f525, plain, (~ (e3 = op(e2, e2)) | ~ (e4 = op(e2, op(e2, e2))) | ~ (e1 = op(op(e2, op(e2, e2)), op(e2, e2))) | ~ (e0 = op(op(e2, e2), op(e2, e2)))), inference(cnf_transformation, [], [f128])).
fof(f128, plain, (~ (e3 = op(e2, e2)) | ~ (e4 = op(e2, op(e2, e2))) | ~ (e1 = op(op(e2, op(e2, e2)), op(e2, e2))) | ~ (e0 = op(op(e2, e2), op(e2, e2)))), inference(ennf_transformation, [], [f8])).
fof(f8, plain, ~ ((e3 = op(e2, e2)) & (e4 = op(e2, op(e2, e2))) & (e1 = op(op(e2, op(e2, e2)), op(e2, e2))) & (e0 = op(op(e2, e2), op(e2, e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax8)).
fof(f1532, plain, (~ spl10_186 | ~ spl10_187 | ~ spl10_188 | ~ spl10_65), inference(avatar_split_clause, [], [f524, f913, f1529, f1525, f1521])).
fof(f524, plain, (~ (e4 = op(e2, e2)) | ~ (e3 = op(e2, op(e2, e2))) | ~ (e1 = op(op(e2, op(e2, e2)), op(e2, e2))) | ~ (e0 = op(op(e2, e2), op(e2, e2)))), inference(cnf_transformation, [], [f127])).
fof(f127, plain, (~ (e4 = op(e2, e2)) | ~ (e3 = op(e2, op(e2, e2))) | ~ (e1 = op(op(e2, op(e2, e2)), op(e2, e2))) | ~ (e0 = op(op(e2, e2), op(e2, e2)))), inference(ennf_transformation, [], [f7])).
fof(f7, plain, ~ ((e4 = op(e2, e2)) & (e3 = op(e2, op(e2, e2))) & (e1 = op(op(e2, op(e2, e2)), op(e2, e2))) & (e0 = op(op(e2, e2), op(e2, e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax7)).
fof(f1519, plain, (spl10_157 | spl10_154 | ~ spl10_183 | ~ spl10_184 | ~ spl10_185), inference(avatar_split_clause, [], [f484, f1513, f1507, f1502, f1349, f1363])).
fof(f1363, plain, (spl10_157 <=> sP0), introduced(avatar_definition, [new_symbols(naming, [spl10_157])])).
fof(f1349, plain, (spl10_154 <=> sP1), introduced(avatar_definition, [new_symbols(naming, [spl10_154])])).
fof(f484, plain, (~ (e0 = op(e4, op(e4, e0))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e0 = op(e2, op(e2, e0))) | sP1 | sP0), inference(cnf_transformation, [], [f257])).
fof(f257, plain, ((((e4 = op(e4, op(e4, e4))) & ~ (e4 = op(e4, op(e4, e4)))) | ((e3 = op(e4, op(e3, e4))) & ~ (e4 = op(e3, op(e3, e4)))) | ((e2 = op(e4, op(e2, e4))) & ~ (e4 = op(e2, op(e2, e4)))) | sP9 | sP8) & (((e4 = op(e3, op(e4, e3))) & ~ (e3 = op(e4, op(e4, e3)))) | ((e3 = op(e3, op(e3, e3))) & ~ (e3 = op(e3, op(e3, e3)))) | ((e2 = op(e3, op(e2, e3))) & ~ (e3 = op(e2, op(e2, e3)))) | sP7 | sP6) & (((e4 = op(e2, op(e4, e2))) & ~ (e2 = op(e4, op(e4, e2)))) | ((e3 = op(e2, op(e3, e2))) & ~ (e2 = op(e3, op(e3, e2)))) | ((e2 = op(e2, op(e2, e2))) & ~ (e2 = op(e2, op(e2, e2)))) | sP5 | sP4) & (((e4 = op(e1, op(e4, e1))) & ~ (e1 = op(e4, op(e4, e1)))) | ((e3 = op(e1, op(e3, e1))) & ~ (e1 = op(e3, op(e3, e1)))) | ((e2 = op(e1, op(e2, e1))) & ~ (e1 = op(e2, op(e2, e1)))) | sP3 | sP2) & (((e4 = op(e0, op(e4, e0))) & ~ (e0 = op(e4, op(e4, e0)))) | ((e3 = op(e0, op(e3, e0))) & ~ (e0 = op(e3, op(e3, e0)))) | ((e2 = op(e0, op(e2, e0))) & ~ (e0 = op(e2, op(e2, e0)))) | sP1 | sP0)), inference(definition_folding, [], [f6, e256, e255, e254, e253, e252, e251, e250, e249, e248, e247])).
fof(f247, plain, (((e0 = op(e0, op(e0, e0))) & ~ (e0 = op(e0, op(e0, e0)))) | ~ sP0), inference(usedef, [], [e247])).
fof(e247, plain, (sP0 <=> ((e0 = op(e0, op(e0, e0))) & ~ (e0 = op(e0, op(e0, e0))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f248, plain, (((e1 = op(e0, op(e1, e0))) & ~ (e0 = op(e1, op(e1, e0)))) | ~ sP1), inference(usedef, [], [e248])).
fof(e248, plain, (sP1 <=> ((e1 = op(e0, op(e1, e0))) & ~ (e0 = op(e1, op(e1, e0))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f249, plain, (((e0 = op(e1, op(e0, e1))) & ~ (e1 = op(e0, op(e0, e1)))) | ~ sP2), inference(usedef, [], [e249])).
fof(e249, plain, (sP2 <=> ((e0 = op(e1, op(e0, e1))) & ~ (e1 = op(e0, op(e0, e1))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f250, plain, (((e1 = op(e1, op(e1, e1))) & ~ (e1 = op(e1, op(e1, e1)))) | ~ sP3), inference(usedef, [], [e250])).
fof(e250, plain, (sP3 <=> ((e1 = op(e1, op(e1, e1))) & ~ (e1 = op(e1, op(e1, e1))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f251, plain, (((e0 = op(e2, op(e0, e2))) & ~ (e2 = op(e0, op(e0, e2)))) | ~ sP4), inference(usedef, [], [e251])).
fof(e251, plain, (sP4 <=> ((e0 = op(e2, op(e0, e2))) & ~ (e2 = op(e0, op(e0, e2))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f252, plain, (((e1 = op(e2, op(e1, e2))) & ~ (e2 = op(e1, op(e1, e2)))) | ~ sP5), inference(usedef, [], [e252])).
fof(e252, plain, (sP5 <=> ((e1 = op(e2, op(e1, e2))) & ~ (e2 = op(e1, op(e1, e2))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f253, plain, (((e0 = op(e3, op(e0, e3))) & ~ (e3 = op(e0, op(e0, e3)))) | ~ sP6), inference(usedef, [], [e253])).
fof(e253, plain, (sP6 <=> ((e0 = op(e3, op(e0, e3))) & ~ (e3 = op(e0, op(e0, e3))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f254, plain, (((e1 = op(e3, op(e1, e3))) & ~ (e3 = op(e1, op(e1, e3)))) | ~ sP7), inference(usedef, [], [e254])).
fof(e254, plain, (sP7 <=> ((e1 = op(e3, op(e1, e3))) & ~ (e3 = op(e1, op(e1, e3))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f255, plain, (((e0 = op(e4, op(e0, e4))) & ~ (e4 = op(e0, op(e0, e4)))) | ~ sP8), inference(usedef, [], [e255])).
fof(e255, plain, (sP8 <=> ((e0 = op(e4, op(e0, e4))) & ~ (e4 = op(e0, op(e0, e4))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f256, plain, (((e1 = op(e4, op(e1, e4))) & ~ (e4 = op(e1, op(e1, e4)))) | ~ sP9), inference(usedef, [], [e256])).
fof(e256, plain, (sP9 <=> ((e1 = op(e4, op(e1, e4))) & ~ (e4 = op(e1, op(e1, e4))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f6, plain, ((((e4 = op(e4, op(e4, e4))) & ~ (e4 = op(e4, op(e4, e4)))) | ((e3 = op(e4, op(e3, e4))) & ~ (e4 = op(e3, op(e3, e4)))) | ((e2 = op(e4, op(e2, e4))) & ~ (e4 = op(e2, op(e2, e4)))) | ((e1 = op(e4, op(e1, e4))) & ~ (e4 = op(e1, op(e1, e4)))) | ((e0 = op(e4, op(e0, e4))) & ~ (e4 = op(e0, op(e0, e4))))) & (((e4 = op(e3, op(e4, e3))) & ~ (e3 = op(e4, op(e4, e3)))) | ((e3 = op(e3, op(e3, e3))) & ~ (e3 = op(e3, op(e3, e3)))) | ((e2 = op(e3, op(e2, e3))) & ~ (e3 = op(e2, op(e2, e3)))) | ((e1 = op(e3, op(e1, e3))) & ~ (e3 = op(e1, op(e1, e3)))) | ((e0 = op(e3, op(e0, e3))) & ~ (e3 = op(e0, op(e0, e3))))) & (((e4 = op(e2, op(e4, e2))) & ~ (e2 = op(e4, op(e4, e2)))) | ((e3 = op(e2, op(e3, e2))) & ~ (e2 = op(e3, op(e3, e2)))) | ((e2 = op(e2, op(e2, e2))) & ~ (e2 = op(e2, op(e2, e2)))) | ((e1 = op(e2, op(e1, e2))) & ~ (e2 = op(e1, op(e1, e2)))) | ((e0 = op(e2, op(e0, e2))) & ~ (e2 = op(e0, op(e0, e2))))) & (((e4 = op(e1, op(e4, e1))) & ~ (e1 = op(e4, op(e4, e1)))) | ((e3 = op(e1, op(e3, e1))) & ~ (e1 = op(e3, op(e3, e1)))) | ((e2 = op(e1, op(e2, e1))) & ~ (e1 = op(e2, op(e2, e1)))) | ((e1 = op(e1, op(e1, e1))) & ~ (e1 = op(e1, op(e1, e1)))) | ((e0 = op(e1, op(e0, e1))) & ~ (e1 = op(e0, op(e0, e1))))) & (((e4 = op(e0, op(e4, e0))) & ~ (e0 = op(e4, op(e4, e0)))) | ((e3 = op(e0, op(e3, e0))) & ~ (e0 = op(e3, op(e3, e0)))) | ((e2 = op(e0, op(e2, e0))) & ~ (e0 = op(e2, op(e2, e0)))) | ((e1 = op(e0, op(e1, e0))) & ~ (e0 = op(e1, op(e1, e0)))) | ((e0 = op(e0, op(e0, e0))) & ~ (e0 = op(e0, op(e0, e0)))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax6)).
fof(f1518, plain, (spl10_157 | spl10_154 | spl10_180 | ~ spl10_184 | ~ spl10_185), inference(avatar_split_clause, [], [f485, f1513, f1507, f1489, f1349, f1363])).
fof(f485, plain, (~ (e0 = op(e4, op(e4, e0))) | ~ (e0 = op(e3, op(e3, e0))) | (e2 = op(e0, op(e2, e0))) | sP1 | sP0), inference(cnf_transformation, [], [f257])).
fof(f1517, plain, (spl10_157 | spl10_154 | ~ spl10_183 | spl10_181 | ~ spl10_185), inference(avatar_split_clause, [], [f486, f1513, f1493, f1502, f1349, f1363])).
fof(f486, plain, (~ (e0 = op(e4, op(e4, e0))) | (e3 = op(e0, op(e3, e0))) | ~ (e0 = op(e2, op(e2, e0))) | sP1 | sP0), inference(cnf_transformation, [], [f257])).
fof(f1516, plain, (spl10_157 | spl10_154 | spl10_180 | spl10_181 | ~ spl10_185), inference(avatar_split_clause, [], [f487, f1513, f1493, f1489, f1349, f1363])).
fof(f487, plain, (~ (e0 = op(e4, op(e4, e0))) | (e3 = op(e0, op(e3, e0))) | (e2 = op(e0, op(e2, e0))) | sP1 | sP0), inference(cnf_transformation, [], [f257])).
fof(f1511, plain, (spl10_157 | spl10_154 | ~ spl10_183 | ~ spl10_184 | spl10_182), inference(avatar_split_clause, [], [f488, f1497, f1507, f1502, f1349, f1363])).
fof(f488, plain, ((e4 = op(e0, op(e4, e0))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e0 = op(e2, op(e2, e0))) | sP1 | sP0), inference(cnf_transformation, [], [f257])).
fof(f1510, plain, (spl10_157 | spl10_154 | spl10_180 | ~ spl10_184 | spl10_182), inference(avatar_split_clause, [], [f489, f1497, f1507, f1489, f1349, f1363])).
fof(f489, plain, ((e4 = op(e0, op(e4, e0))) | ~ (e0 = op(e3, op(e3, e0))) | (e2 = op(e0, op(e2, e0))) | sP1 | sP0), inference(cnf_transformation, [], [f257])).
fof(f1505, plain, (spl10_157 | spl10_154 | ~ spl10_183 | spl10_181 | spl10_182), inference(avatar_split_clause, [], [f490, f1497, f1493, f1502, f1349, f1363])).
fof(f490, plain, ((e4 = op(e0, op(e4, e0))) | (e3 = op(e0, op(e3, e0))) | ~ (e0 = op(e2, op(e2, e0))) | sP1 | sP0), inference(cnf_transformation, [], [f257])).
fof(f1500, plain, (spl10_157 | spl10_154 | spl10_180 | spl10_181 | spl10_182), inference(avatar_split_clause, [], [f491, f1497, f1493, f1489, f1349, f1363])).
fof(f491, plain, ((e4 = op(e0, op(e4, e0))) | (e3 = op(e0, op(e3, e0))) | (e2 = op(e0, op(e2, e0))) | sP1 | sP0), inference(cnf_transformation, [], [f257])).
fof(f1487, plain, (spl10_151 | spl10_149 | ~ spl10_177 | ~ spl10_178 | ~ spl10_179), inference(avatar_split_clause, [], [f492, f1481, f1475, f1470, f1325, f1335])).
fof(f1335, plain, (spl10_151 <=> sP2), introduced(avatar_definition, [new_symbols(naming, [spl10_151])])).
fof(f1325, plain, (spl10_149 <=> sP3), introduced(avatar_definition, [new_symbols(naming, [spl10_149])])).
fof(f492, plain, (~ (e1 = op(e4, op(e4, e1))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e1 = op(e2, op(e2, e1))) | sP3 | sP2), inference(cnf_transformation, [], [f257])).
fof(f1486, plain, (spl10_151 | spl10_149 | spl10_174 | ~ spl10_178 | ~ spl10_179), inference(avatar_split_clause, [], [f493, f1481, f1475, f1457, f1325, f1335])).
fof(f493, plain, (~ (e1 = op(e4, op(e4, e1))) | ~ (e1 = op(e3, op(e3, e1))) | (e2 = op(e1, op(e2, e1))) | sP3 | sP2), inference(cnf_transformation, [], [f257])).
fof(f1485, plain, (spl10_151 | spl10_149 | ~ spl10_177 | spl10_175 | ~ spl10_179), inference(avatar_split_clause, [], [f494, f1481, f1461, f1470, f1325, f1335])).
fof(f494, plain, (~ (e1 = op(e4, op(e4, e1))) | (e3 = op(e1, op(e3, e1))) | ~ (e1 = op(e2, op(e2, e1))) | sP3 | sP2), inference(cnf_transformation, [], [f257])).
fof(f1484, plain, (spl10_151 | spl10_149 | spl10_174 | spl10_175 | ~ spl10_179), inference(avatar_split_clause, [], [f495, f1481, f1461, f1457, f1325, f1335])).
fof(f495, plain, (~ (e1 = op(e4, op(e4, e1))) | (e3 = op(e1, op(e3, e1))) | (e2 = op(e1, op(e2, e1))) | sP3 | sP2), inference(cnf_transformation, [], [f257])).
fof(f1479, plain, (spl10_151 | spl10_149 | ~ spl10_177 | ~ spl10_178 | spl10_176), inference(avatar_split_clause, [], [f496, f1465, f1475, f1470, f1325, f1335])).
fof(f496, plain, ((e4 = op(e1, op(e4, e1))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e1 = op(e2, op(e2, e1))) | sP3 | sP2), inference(cnf_transformation, [], [f257])).
fof(f1478, plain, (spl10_151 | spl10_149 | spl10_174 | ~ spl10_178 | spl10_176), inference(avatar_split_clause, [], [f497, f1465, f1475, f1457, f1325, f1335])).
fof(f497, plain, ((e4 = op(e1, op(e4, e1))) | ~ (e1 = op(e3, op(e3, e1))) | (e2 = op(e1, op(e2, e1))) | sP3 | sP2), inference(cnf_transformation, [], [f257])).
fof(f1473, plain, (spl10_151 | spl10_149 | ~ spl10_177 | spl10_175 | spl10_176), inference(avatar_split_clause, [], [f498, f1465, f1461, f1470, f1325, f1335])).
fof(f498, plain, ((e4 = op(e1, op(e4, e1))) | (e3 = op(e1, op(e3, e1))) | ~ (e1 = op(e2, op(e2, e1))) | sP3 | sP2), inference(cnf_transformation, [], [f257])).
fof(f1468, plain, (spl10_151 | spl10_149 | spl10_174 | spl10_175 | spl10_176), inference(avatar_split_clause, [], [f499, f1465, f1461, f1457, f1325, f1335])).
fof(f499, plain, ((e4 = op(e1, op(e4, e1))) | (e3 = op(e1, op(e3, e1))) | (e2 = op(e1, op(e2, e1))) | sP3 | sP2), inference(cnf_transformation, [], [f257])).
fof(f1455, plain, (spl10_146 | spl10_143 | ~ spl10_169 | ~ spl10_172 | ~ spl10_173), inference(avatar_split_clause, [], [f500, f1449, f1443, f1429, f1297, f1311])).
fof(f1311, plain, (spl10_146 <=> sP4), introduced(avatar_definition, [new_symbols(naming, [spl10_146])])).
fof(f1297, plain, (spl10_143 <=> sP5), introduced(avatar_definition, [new_symbols(naming, [spl10_143])])).
fof(f500, plain, (~ (e2 = op(e4, op(e4, e2))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e2 = op(e2, op(e2, e2))) | sP5 | sP4), inference(cnf_transformation, [], [f257])).
fof(f1454, plain, (spl10_146 | spl10_143 | spl10_169 | ~ spl10_172 | ~ spl10_173), inference(avatar_split_clause, [], [f501, f1449, f1443, f1429, f1297, f1311])).
fof(f501, plain, (~ (e2 = op(e4, op(e4, e2))) | ~ (e2 = op(e3, op(e3, e2))) | (e2 = op(e2, op(e2, e2))) | sP5 | sP4), inference(cnf_transformation, [], [f257])).
fof(f1453, plain, (spl10_146 | spl10_143 | ~ spl10_169 | spl10_170 | ~ spl10_173), inference(avatar_split_clause, [], [f502, f1449, f1433, f1429, f1297, f1311])).
fof(f502, plain, (~ (e2 = op(e4, op(e4, e2))) | (e3 = op(e2, op(e3, e2))) | ~ (e2 = op(e2, op(e2, e2))) | sP5 | sP4), inference(cnf_transformation, [], [f257])).
fof(f1452, plain, (spl10_146 | spl10_143 | spl10_169 | spl10_170 | ~ spl10_173), inference(avatar_split_clause, [], [f503, f1449, f1433, f1429, f1297, f1311])).
fof(f503, plain, (~ (e2 = op(e4, op(e4, e2))) | (e3 = op(e2, op(e3, e2))) | (e2 = op(e2, op(e2, e2))) | sP5 | sP4), inference(cnf_transformation, [], [f257])).
fof(f1447, plain, (spl10_146 | spl10_143 | ~ spl10_169 | ~ spl10_172 | spl10_171), inference(avatar_split_clause, [], [f504, f1437, f1443, f1429, f1297, f1311])).
fof(f504, plain, ((e4 = op(e2, op(e4, e2))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e2 = op(e2, op(e2, e2))) | sP5 | sP4), inference(cnf_transformation, [], [f257])).
fof(f1446, plain, (spl10_146 | spl10_143 | spl10_169 | ~ spl10_172 | spl10_171), inference(avatar_split_clause, [], [f505, f1437, f1443, f1429, f1297, f1311])).
fof(f505, plain, ((e4 = op(e2, op(e4, e2))) | ~ (e2 = op(e3, op(e3, e2))) | (e2 = op(e2, op(e2, e2))) | sP5 | sP4), inference(cnf_transformation, [], [f257])).
fof(f1441, plain, (spl10_146 | spl10_143 | ~ spl10_169 | spl10_170 | spl10_171), inference(avatar_split_clause, [], [f506, f1437, f1433, f1429, f1297, f1311])).
fof(f506, plain, ((e4 = op(e2, op(e4, e2))) | (e3 = op(e2, op(e3, e2))) | ~ (e2 = op(e2, op(e2, e2))) | sP5 | sP4), inference(cnf_transformation, [], [f257])).
fof(f1440, plain, (spl10_146 | spl10_143 | spl10_169 | spl10_170 | spl10_171), inference(avatar_split_clause, [], [f507, f1437, f1433, f1429, f1297, f1311])).
fof(f507, plain, ((e4 = op(e2, op(e4, e2))) | (e3 = op(e2, op(e3, e2))) | (e2 = op(e2, op(e2, e2))) | sP5 | sP4), inference(cnf_transformation, [], [f257])).
fof(f1427, plain, (spl10_140 | spl10_137 | ~ spl10_167 | ~ spl10_165 | ~ spl10_168), inference(avatar_split_clause, [], [f508, f1421, f1405, f1414, f1269, f1283])).
fof(f1283, plain, (spl10_140 <=> sP6), introduced(avatar_definition, [new_symbols(naming, [spl10_140])])).
fof(f1269, plain, (spl10_137 <=> sP7), introduced(avatar_definition, [new_symbols(naming, [spl10_137])])).
fof(f508, plain, (~ (e3 = op(e4, op(e4, e3))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e3 = op(e2, op(e2, e3))) | sP7 | sP6), inference(cnf_transformation, [], [f257])).
fof(f1426, plain, (spl10_140 | spl10_137 | spl10_164 | ~ spl10_165 | ~ spl10_168), inference(avatar_split_clause, [], [f509, f1421, f1405, f1401, f1269, f1283])).
fof(f509, plain, (~ (e3 = op(e4, op(e4, e3))) | ~ (e3 = op(e3, op(e3, e3))) | (e2 = op(e3, op(e2, e3))) | sP7 | sP6), inference(cnf_transformation, [], [f257])).
fof(f1425, plain, (spl10_140 | spl10_137 | ~ spl10_167 | spl10_165 | ~ spl10_168), inference(avatar_split_clause, [], [f510, f1421, f1405, f1414, f1269, f1283])).
fof(f510, plain, (~ (e3 = op(e4, op(e4, e3))) | (e3 = op(e3, op(e3, e3))) | ~ (e3 = op(e2, op(e2, e3))) | sP7 | sP6), inference(cnf_transformation, [], [f257])).
fof(f1424, plain, (spl10_140 | spl10_137 | spl10_164 | spl10_165 | ~ spl10_168), inference(avatar_split_clause, [], [f511, f1421, f1405, f1401, f1269, f1283])).
fof(f511, plain, (~ (e3 = op(e4, op(e4, e3))) | (e3 = op(e3, op(e3, e3))) | (e2 = op(e3, op(e2, e3))) | sP7 | sP6), inference(cnf_transformation, [], [f257])).
fof(f1419, plain, (spl10_140 | spl10_137 | ~ spl10_167 | ~ spl10_165 | spl10_166), inference(avatar_split_clause, [], [f512, f1409, f1405, f1414, f1269, f1283])).
fof(f512, plain, ((e4 = op(e3, op(e4, e3))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e3 = op(e2, op(e2, e3))) | sP7 | sP6), inference(cnf_transformation, [], [f257])).
fof(f1418, plain, (spl10_140 | spl10_137 | spl10_164 | ~ spl10_165 | spl10_166), inference(avatar_split_clause, [], [f513, f1409, f1405, f1401, f1269, f1283])).
fof(f513, plain, ((e4 = op(e3, op(e4, e3))) | ~ (e3 = op(e3, op(e3, e3))) | (e2 = op(e3, op(e2, e3))) | sP7 | sP6), inference(cnf_transformation, [], [f257])).
fof(f1417, plain, (spl10_140 | spl10_137 | ~ spl10_167 | spl10_165 | spl10_166), inference(avatar_split_clause, [], [f514, f1409, f1405, f1414, f1269, f1283])).
fof(f514, plain, ((e4 = op(e3, op(e4, e3))) | (e3 = op(e3, op(e3, e3))) | ~ (e3 = op(e2, op(e2, e3))) | sP7 | sP6), inference(cnf_transformation, [], [f257])).
fof(f1412, plain, (spl10_140 | spl10_137 | spl10_164 | spl10_165 | spl10_166), inference(avatar_split_clause, [], [f515, f1409, f1405, f1401, f1269, f1283])).
fof(f515, plain, ((e4 = op(e3, op(e4, e3))) | (e3 = op(e3, op(e3, e3))) | (e2 = op(e3, op(e2, e3))) | sP7 | sP6), inference(cnf_transformation, [], [f257])).
fof(f1399, plain, (spl10_134 | spl10_131 | ~ spl10_162 | ~ spl10_163 | ~ spl10_161), inference(avatar_split_clause, [], [f516, f1381, f1391, f1386, f1241, f1255])).
fof(f1255, plain, (spl10_134 <=> sP8), introduced(avatar_definition, [new_symbols(naming, [spl10_134])])).
fof(f1241, plain, (spl10_131 <=> sP9), introduced(avatar_definition, [new_symbols(naming, [spl10_131])])).
fof(f516, plain, (~ (e4 = op(e4, op(e4, e4))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e4 = op(e2, op(e2, e4))) | sP9 | sP8), inference(cnf_transformation, [], [f257])).
fof(f1398, plain, (spl10_134 | spl10_131 | spl10_159 | ~ spl10_163 | ~ spl10_161), inference(avatar_split_clause, [], [f517, f1381, f1391, f1373, f1241, f1255])).
fof(f517, plain, (~ (e4 = op(e4, op(e4, e4))) | ~ (e4 = op(e3, op(e3, e4))) | (e2 = op(e4, op(e2, e4))) | sP9 | sP8), inference(cnf_transformation, [], [f257])).
fof(f1397, plain, (spl10_134 | spl10_131 | ~ spl10_162 | spl10_160 | ~ spl10_161), inference(avatar_split_clause, [], [f518, f1381, f1377, f1386, f1241, f1255])).
fof(f518, plain, (~ (e4 = op(e4, op(e4, e4))) | (e3 = op(e4, op(e3, e4))) | ~ (e4 = op(e2, op(e2, e4))) | sP9 | sP8), inference(cnf_transformation, [], [f257])).
fof(f1396, plain, (spl10_134 | spl10_131 | spl10_159 | spl10_160 | ~ spl10_161), inference(avatar_split_clause, [], [f519, f1381, f1377, f1373, f1241, f1255])).
fof(f519, plain, (~ (e4 = op(e4, op(e4, e4))) | (e3 = op(e4, op(e3, e4))) | (e2 = op(e4, op(e2, e4))) | sP9 | sP8), inference(cnf_transformation, [], [f257])).
fof(f1395, plain, (spl10_134 | spl10_131 | ~ spl10_162 | ~ spl10_163 | spl10_161), inference(avatar_split_clause, [], [f520, f1381, f1391, f1386, f1241, f1255])).
fof(f520, plain, ((e4 = op(e4, op(e4, e4))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e4 = op(e2, op(e2, e4))) | sP9 | sP8), inference(cnf_transformation, [], [f257])).
fof(f1394, plain, (spl10_134 | spl10_131 | spl10_159 | ~ spl10_163 | spl10_161), inference(avatar_split_clause, [], [f521, f1381, f1391, f1373, f1241, f1255])).
fof(f521, plain, ((e4 = op(e4, op(e4, e4))) | ~ (e4 = op(e3, op(e3, e4))) | (e2 = op(e4, op(e2, e4))) | sP9 | sP8), inference(cnf_transformation, [], [f257])).
fof(f1389, plain, (spl10_134 | spl10_131 | ~ spl10_162 | spl10_160 | spl10_161), inference(avatar_split_clause, [], [f522, f1381, f1377, f1386, f1241, f1255])).
fof(f522, plain, ((e4 = op(e4, op(e4, e4))) | (e3 = op(e4, op(e3, e4))) | ~ (e4 = op(e2, op(e2, e4))) | sP9 | sP8), inference(cnf_transformation, [], [f257])).
fof(f1384, plain, (spl10_134 | spl10_131 | spl10_159 | spl10_160 | spl10_161), inference(avatar_split_clause, [], [f523, f1381, f1377, f1373, f1241, f1255])).
fof(f523, plain, ((e4 = op(e4, op(e4, e4))) | (e3 = op(e4, op(e3, e4))) | (e2 = op(e4, op(e2, e4))) | sP9 | sP8), inference(cnf_transformation, [], [f257])).
fof(f1371, plain, (~ spl10_157 | ~ spl10_158), inference(avatar_split_clause, [], [f482, f1367, f1363])).
fof(f1367, plain, (spl10_158 <=> (e0 = op(e0, op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl10_158])])).
fof(f482, plain, (~ (e0 = op(e0, op(e0, e0))) | ~ sP0), inference(cnf_transformation, [], [f267])).
fof(f267, plain, (((e0 = op(e0, op(e0, e0))) & ~ (e0 = op(e0, op(e0, e0)))) | ~ sP0), inference(nnf_transformation, [], [f247])).
fof(f1370, plain, (~ spl10_157 | spl10_158), inference(avatar_split_clause, [], [f483, f1367, f1363])).
fof(f483, plain, ((e0 = op(e0, op(e0, e0))) | ~ sP0), inference(cnf_transformation, [], [f267])).
fof(f1361, plain, (~ spl10_154 | ~ spl10_156), inference(avatar_split_clause, [], [f480, f1358, f1349])).
fof(f480, plain, (~ (e0 = op(e1, op(e1, e0))) | ~ sP1), inference(cnf_transformation, [], [f266])).
fof(f266, plain, (((e1 = op(e0, op(e1, e0))) & ~ (e0 = op(e1, op(e1, e0)))) | ~ sP1), inference(nnf_transformation, [], [f248])).
fof(f1356, plain, (~ spl10_154 | spl10_155), inference(avatar_split_clause, [], [f481, f1353, f1349])).
fof(f481, plain, ((e1 = op(e0, op(e1, e0))) | ~ sP1), inference(cnf_transformation, [], [f266])).
fof(f1347, plain, (~ spl10_151 | ~ spl10_153), inference(avatar_split_clause, [], [f478, f1344, f1335])).
fof(f478, plain, (~ (e1 = op(e0, op(e0, e1))) | ~ sP2), inference(cnf_transformation, [], [f265])).
fof(f265, plain, (((e0 = op(e1, op(e0, e1))) & ~ (e1 = op(e0, op(e0, e1)))) | ~ sP2), inference(nnf_transformation, [], [f249])).
fof(f1342, plain, (~ spl10_151 | spl10_152), inference(avatar_split_clause, [], [f479, f1339, f1335])).
fof(f479, plain, ((e0 = op(e1, op(e0, e1))) | ~ sP2), inference(cnf_transformation, [], [f265])).
fof(f1333, plain, (~ spl10_149 | ~ spl10_150), inference(avatar_split_clause, [], [f476, f1329, f1325])).
fof(f1329, plain, (spl10_150 <=> (e1 = op(e1, op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl10_150])])).
fof(f476, plain, (~ (e1 = op(e1, op(e1, e1))) | ~ sP3), inference(cnf_transformation, [], [f264])).
fof(f264, plain, (((e1 = op(e1, op(e1, e1))) & ~ (e1 = op(e1, op(e1, e1)))) | ~ sP3), inference(nnf_transformation, [], [f250])).
fof(f1332, plain, (~ spl10_149 | spl10_150), inference(avatar_split_clause, [], [f477, f1329, f1325])).
fof(f477, plain, ((e1 = op(e1, op(e1, e1))) | ~ sP3), inference(cnf_transformation, [], [f264])).
fof(f1323, plain, (~ spl10_146 | ~ spl10_148), inference(avatar_split_clause, [], [f474, f1320, f1311])).
fof(f474, plain, (~ (e2 = op(e0, op(e0, e2))) | ~ sP4), inference(cnf_transformation, [], [f263])).
fof(f263, plain, (((e0 = op(e2, op(e0, e2))) & ~ (e2 = op(e0, op(e0, e2)))) | ~ sP4), inference(nnf_transformation, [], [f251])).
fof(f1318, plain, (~ spl10_146 | spl10_147), inference(avatar_split_clause, [], [f475, f1315, f1311])).
fof(f475, plain, ((e0 = op(e2, op(e0, e2))) | ~ sP4), inference(cnf_transformation, [], [f263])).
fof(f1309, plain, (~ spl10_143 | ~ spl10_145), inference(avatar_split_clause, [], [f472, f1306, f1297])).
fof(f472, plain, (~ (e2 = op(e1, op(e1, e2))) | ~ sP5), inference(cnf_transformation, [], [f262])).
fof(f262, plain, (((e1 = op(e2, op(e1, e2))) & ~ (e2 = op(e1, op(e1, e2)))) | ~ sP5), inference(nnf_transformation, [], [f252])).
fof(f1304, plain, (~ spl10_143 | spl10_144), inference(avatar_split_clause, [], [f473, f1301, f1297])).
fof(f473, plain, ((e1 = op(e2, op(e1, e2))) | ~ sP5), inference(cnf_transformation, [], [f262])).
fof(f1295, plain, (~ spl10_140 | ~ spl10_142), inference(avatar_split_clause, [], [f470, f1292, f1283])).
fof(f470, plain, (~ (e3 = op(e0, op(e0, e3))) | ~ sP6), inference(cnf_transformation, [], [f261])).
fof(f261, plain, (((e0 = op(e3, op(e0, e3))) & ~ (e3 = op(e0, op(e0, e3)))) | ~ sP6), inference(nnf_transformation, [], [f253])).
fof(f1290, plain, (~ spl10_140 | spl10_141), inference(avatar_split_clause, [], [f471, f1287, f1283])).
fof(f471, plain, ((e0 = op(e3, op(e0, e3))) | ~ sP6), inference(cnf_transformation, [], [f261])).
fof(f1281, plain, (~ spl10_137 | ~ spl10_139), inference(avatar_split_clause, [], [f468, f1278, f1269])).
fof(f468, plain, (~ (e3 = op(e1, op(e1, e3))) | ~ sP7), inference(cnf_transformation, [], [f260])).
fof(f260, plain, (((e1 = op(e3, op(e1, e3))) & ~ (e3 = op(e1, op(e1, e3)))) | ~ sP7), inference(nnf_transformation, [], [f254])).
fof(f1276, plain, (~ spl10_137 | spl10_138), inference(avatar_split_clause, [], [f469, f1273, f1269])).
fof(f469, plain, ((e1 = op(e3, op(e1, e3))) | ~ sP7), inference(cnf_transformation, [], [f260])).
fof(f1267, plain, (~ spl10_134 | ~ spl10_136), inference(avatar_split_clause, [], [f466, f1264, f1255])).
fof(f466, plain, (~ (e4 = op(e0, op(e0, e4))) | ~ sP8), inference(cnf_transformation, [], [f259])).
fof(f259, plain, (((e0 = op(e4, op(e0, e4))) & ~ (e4 = op(e0, op(e0, e4)))) | ~ sP8), inference(nnf_transformation, [], [f255])).
fof(f1262, plain, (~ spl10_134 | spl10_135), inference(avatar_split_clause, [], [f467, f1259, f1255])).
fof(f467, plain, ((e0 = op(e4, op(e0, e4))) | ~ sP8), inference(cnf_transformation, [], [f259])).
fof(f1253, plain, (~ spl10_131 | ~ spl10_133), inference(avatar_split_clause, [], [f464, f1250, f1241])).
fof(f464, plain, (~ (e4 = op(e1, op(e1, e4))) | ~ sP9), inference(cnf_transformation, [], [f258])).
fof(f258, plain, (((e1 = op(e4, op(e1, e4))) & ~ (e4 = op(e1, op(e1, e4)))) | ~ sP9), inference(nnf_transformation, [], [f256])).
fof(f1248, plain, (~ spl10_131 | spl10_132), inference(avatar_split_clause, [], [f465, f1245, f1241])).
fof(f465, plain, ((e1 = op(e4, op(e1, e4))) | ~ sP9), inference(cnf_transformation, [], [f258])).
fof(f1239, plain, (spl10_121 | spl10_116 | spl10_111 | spl10_106 | spl10_101), inference(avatar_split_clause, [], [f304, f1065, f1086, f1107, f1128, f1149])).
fof(f304, plain, ((e0 = op(e0, e4)) | (e0 = op(e0, e3)) | (e0 = op(e0, e2)) | (e0 = op(e0, e1)) | (e0 = op(e0, e0))), inference(cnf_transformation, [], [f3])).
fof(f3, plain, (((e4 = op(e4, e4)) | (e4 = op(e3, e4)) | (e4 = op(e2, e4)) | (e4 = op(e1, e4)) | (e4 = op(e0, e4))) & ((e4 = op(e4, e4)) | (e4 = op(e4, e3)) | (e4 = op(e4, e2)) | (e4 = op(e4, e1)) | (e4 = op(e4, e0))) & ((e3 = op(e4, e4)) | (e3 = op(e3, e4)) | (e3 = op(e2, e4)) | (e3 = op(e1, e4)) | (e3 = op(e0, e4))) & ((e3 = op(e4, e4)) | (e3 = op(e4, e3)) | (e3 = op(e4, e2)) | (e3 = op(e4, e1)) | (e3 = op(e4, e0))) & ((e2 = op(e4, e4)) | (e2 = op(e3, e4)) | (e2 = op(e2, e4)) | (e2 = op(e1, e4)) | (e2 = op(e0, e4))) & ((e2 = op(e4, e4)) | (e2 = op(e4, e3)) | (e2 = op(e4, e2)) | (e2 = op(e4, e1)) | (e2 = op(e4, e0))) & ((e1 = op(e4, e4)) | (e1 = op(e3, e4)) | (e1 = op(e2, e4)) | (e1 = op(e1, e4)) | (e1 = op(e0, e4))) & ((e1 = op(e4, e4)) | (e1 = op(e4, e3)) | (e1 = op(e4, e2)) | (e1 = op(e4, e1)) | (e1 = op(e4, e0))) & ((e0 = op(e4, e4)) | (e0 = op(e3, e4)) | (e0 = op(e2, e4)) | (e0 = op(e1, e4)) | (e0 = op(e0, e4))) & ((e0 = op(e4, e4)) | (e0 = op(e4, e3)) | (e0 = op(e4, e2)) | (e0 = op(e4, e1)) | (e0 = op(e4, e0))) & ((e4 = op(e4, e3)) | (e4 = op(e3, e3)) | (e4 = op(e2, e3)) | (e4 = op(e1, e3)) | (e4 = op(e0, e3))) & ((e4 = op(e3, e4)) | (e4 = op(e3, e3)) | (e4 = op(e3, e2)) | (e4 = op(e3, e1)) | (e4 = op(e3, e0))) & ((e3 = op(e4, e3)) | (e3 = op(e3, e3)) | (e3 = op(e2, e3)) | (e3 = op(e1, e3)) | (e3 = op(e0, e3))) & ((e3 = op(e3, e4)) | (e3 = op(e3, e3)) | (e3 = op(e3, e2)) | (e3 = op(e3, e1)) | (e3 = op(e3, e0))) & ((e2 = op(e4, e3)) | (e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(e0, e3))) & ((e2 = op(e3, e4)) | (e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, e0))) & ((e1 = op(e4, e3)) | (e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(e0, e3))) & ((e1 = op(e3, e4)) | (e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, e0))) & ((e0 = op(e4, e3)) | (e0 = op(e3, e3)) | (e0 = op(e2, e3)) | (e0 = op(e1, e3)) | (e0 = op(e0, e3))) & ((e0 = op(e3, e4)) | (e0 = op(e3, e3)) | (e0 = op(e3, e2)) | (e0 = op(e3, e1)) | (e0 = op(e3, e0))) & ((e4 = op(e4, e2)) | (e4 = op(e3, e2)) | (e4 = op(e2, e2)) | (e4 = op(e1, e2)) | (e4 = op(e0, e2))) & ((e4 = op(e2, e4)) | (e4 = op(e2, e3)) | (e4 = op(e2, e2)) | (e4 = op(e2, e1)) | (e4 = op(e2, e0))) & ((e3 = op(e4, e2)) | (e3 = op(e3, e2)) | (e3 = op(e2, e2)) | (e3 = op(e1, e2)) | (e3 = op(e0, e2))) & ((e3 = op(e2, e4)) | (e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, e0))) & ((e2 = op(e4, e2)) | (e2 = op(e3, e2)) | (e2 = op(e2, e2)) | (e2 = op(e1, e2)) | (e2 = op(e0, e2))) & ((e2 = op(e2, e4)) | (e2 = op(e2, e3)) | (e2 = op(e2, e2)) | (e2 = op(e2, e1)) | (e2 = op(e2, e0))) & ((e1 = op(e4, e2)) | (e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(e0, e2))) & ((e1 = op(e2, e4)) | (e1 = op(e2, e3)) | (e1 = op(e2, e2)) | (e1 = op(e2, e1)) | (e1 = op(e2, e0))) & ((e0 = op(e4, e2)) | (e0 = op(e3, e2)) | (e0 = op(e2, e2)) | (e0 = op(e1, e2)) | (e0 = op(e0, e2))) & ((e0 = op(e2, e4)) | (e0 = op(e2, e3)) | (e0 = op(e2, e2)) | (e0 = op(e2, e1)) | (e0 = op(e2, e0))) & ((e4 = op(e4, e1)) | (e4 = op(e3, e1)) | (e4 = op(e2, e1)) | (e4 = op(e1, e1)) | (e4 = op(e0, e1))) & ((e4 = op(e1, e4)) | (e4 = op(e1, e3)) | (e4 = op(e1, e2)) | (e4 = op(e1, e1)) | (e4 = op(e1, e0))) & ((e3 = op(e4, e1)) | (e3 = op(e3, e1)) | (e3 = op(e2, e1)) | (e3 = op(e1, e1)) | (e3 = op(e0, e1))) & ((e3 = op(e1, e4)) | (e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, e0))) & ((e2 = op(e4, e1)) | (e2 = op(e3, e1)) | (e2 = op(e2, e1)) | (e2 = op(e1, e1)) | (e2 = op(e0, e1))) & ((e2 = op(e1, e4)) | (e2 = op(e1, e3)) | (e2 = op(e1, e2)) | (e2 = op(e1, e1)) | (e2 = op(e1, e0))) & ((e1 = op(e4, e1)) | (e1 = op(e3, e1)) | (e1 = op(e2, e1)) | (e1 = op(e1, e1)) | (e1 = op(e0, e1))) & ((e1 = op(e1, e4)) | (e1 = op(e1, e3)) | (e1 = op(e1, e2)) | (e1 = op(e1, e1)) | (e1 = op(e1, e0))) & ((e0 = op(e4, e1)) | (e0 = op(e3, e1)) | (e0 = op(e2, e1)) | (e0 = op(e1, e1)) | (e0 = op(e0, e1))) & ((e0 = op(e1, e4)) | (e0 = op(e1, e3)) | (e0 = op(e1, e2)) | (e0 = op(e1, e1)) | (e0 = op(e1, e0))) & ((e4 = op(e4, e0)) | (e4 = op(e3, e0)) | (e4 = op(e2, e0)) | (e4 = op(e1, e0)) | (op(e0, e0) = e4)) & ((e4 = op(e0, e4)) | (e4 = op(e0, e3)) | (e4 = op(e0, e2)) | (e4 = op(e0, e1)) | (op(e0, e0) = e4)) & ((e3 = op(e4, e0)) | (e3 = op(e3, e0)) | (e3 = op(e2, e0)) | (e3 = op(e1, e0)) | (op(e0, e0) = e3)) & ((e3 = op(e0, e4)) | (e3 = op(e0, e3)) | (e3 = op(e0, e2)) | (e3 = op(e0, e1)) | (op(e0, e0) = e3)) & ((e2 = op(e4, e0)) | (e2 = op(e3, e0)) | (e2 = op(e2, e0)) | (e2 = op(e1, e0)) | (op(e0, e0) = e2)) & ((e2 = op(e0, e4)) | (e2 = op(e0, e3)) | (e2 = op(e0, e2)) | (e2 = op(e0, e1)) | (op(e0, e0) = e2)) & ((e1 = op(e4, e0)) | (e1 = op(e3, e0)) | (e1 = op(e2, e0)) | (e1 = op(e1, e0)) | (op(e0, e0) = e1)) & ((e1 = op(e0, e4)) | (e1 = op(e0, e3)) | (e1 = op(e0, e2)) | (e1 = op(e0, e1)) | (op(e0, e0) = e1)) & ((e0 = op(e4, e0)) | (e0 = op(e3, e0)) | (e0 = op(e2, e0)) | (e0 = op(e1, e0)) | (e0 = op(e0, e0))) & ((e0 = op(e0, e4)) | (e0 = op(e0, e3)) | (e0 = op(e0, e2)) | (e0 = op(e0, e1)) | (e0 = op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax3)).
fof(f1237, plain, (spl10_122 | spl10_117 | spl10_112 | spl10_107 | spl10_102), inference(avatar_split_clause, [], [f306, f1069, f1090, f1111, f1132, f1153])).
fof(f306, plain, ((e1 = op(e0, e4)) | (e1 = op(e0, e3)) | (e1 = op(e0, e2)) | (e1 = op(e0, e1)) | (op(e0, e0) = e1)), inference(cnf_transformation, [], [f3])).
fof(f1236, plain, (spl10_122 | spl10_97 | spl10_72 | spl10_47 | spl10_22), inference(avatar_split_clause, [], [f307, f733, f838, f943, f1048, f1153])).
fof(f307, plain, ((e1 = op(e4, e0)) | (e1 = op(e3, e0)) | (e1 = op(e2, e0)) | (e1 = op(e1, e0)) | (op(e0, e0) = e1)), inference(cnf_transformation, [], [f3])).
fof(f1235, plain, (spl10_123 | spl10_118 | spl10_113 | spl10_108 | spl10_103), inference(avatar_split_clause, [], [f308, f1073, f1094, f1115, f1136, f1157])).
fof(f308, plain, ((e2 = op(e0, e4)) | (e2 = op(e0, e3)) | (e2 = op(e0, e2)) | (e2 = op(e0, e1)) | (op(e0, e0) = e2)), inference(cnf_transformation, [], [f3])).
fof(f1234, plain, (spl10_123 | spl10_98 | spl10_73 | spl10_48 | spl10_23), inference(avatar_split_clause, [], [f309, f737, f842, f947, f1052, f1157])).
fof(f309, plain, ((e2 = op(e4, e0)) | (e2 = op(e3, e0)) | (e2 = op(e2, e0)) | (e2 = op(e1, e0)) | (op(e0, e0) = e2)), inference(cnf_transformation, [], [f3])).
fof(f1233, plain, (spl10_124 | spl10_119 | spl10_114 | spl10_109 | spl10_104), inference(avatar_split_clause, [], [f310, f1077, f1098, f1119, f1140, f1161])).
fof(f310, plain, ((e3 = op(e0, e4)) | (e3 = op(e0, e3)) | (e3 = op(e0, e2)) | (e3 = op(e0, e1)) | (op(e0, e0) = e3)), inference(cnf_transformation, [], [f3])).
fof(f1232, plain, (spl10_124 | spl10_99 | spl10_74 | spl10_49 | spl10_24), inference(avatar_split_clause, [], [f311, f741, f846, f951, f1056, f1161])).
fof(f311, plain, ((e3 = op(e4, e0)) | (e3 = op(e3, e0)) | (e3 = op(e2, e0)) | (e3 = op(e1, e0)) | (op(e0, e0) = e3)), inference(cnf_transformation, [], [f3])).
fof(f1231, plain, (spl10_125 | spl10_120 | spl10_115 | spl10_110 | spl10_105), inference(avatar_split_clause, [], [f312, f1081, f1102, f1123, f1144, f1165])).
fof(f312, plain, ((e4 = op(e0, e4)) | (e4 = op(e0, e3)) | (e4 = op(e0, e2)) | (e4 = op(e0, e1)) | (op(e0, e0) = e4)), inference(cnf_transformation, [], [f3])).
fof(f1230, plain, (spl10_125 | spl10_100 | spl10_75 | spl10_50 | spl10_25), inference(avatar_split_clause, [], [f313, f745, f850, f955, f1060, f1165])).
fof(f313, plain, ((e4 = op(e4, e0)) | (e4 = op(e3, e0)) | (e4 = op(e2, e0)) | (e4 = op(e1, e0)) | (op(e0, e0) = e4)), inference(cnf_transformation, [], [f3])).
fof(f1229, plain, (spl10_96 | spl10_91 | spl10_86 | spl10_81 | spl10_76), inference(avatar_split_clause, [], [f314, f960, f981, f1002, f1023, f1044])).
fof(f314, plain, ((e0 = op(e1, e4)) | (e0 = op(e1, e3)) | (e0 = op(e1, e2)) | (e0 = op(e1, e1)) | (e0 = op(e1, e0))), inference(cnf_transformation, [], [f3])).
fof(f1228, plain, (spl10_116 | spl10_91 | spl10_66 | spl10_41 | spl10_16), inference(avatar_split_clause, [], [f315, f708, f813, f918, f1023, f1128])).
fof(f315, plain, ((e0 = op(e4, e1)) | (e0 = op(e3, e1)) | (e0 = op(e2, e1)) | (e0 = op(e1, e1)) | (e0 = op(e0, e1))), inference(cnf_transformation, [], [f3])).
fof(f1225, plain, (spl10_98 | spl10_93 | spl10_88 | spl10_83 | spl10_78), inference(avatar_split_clause, [], [f318, f968, f989, f1010, f1031, f1052])).
fof(f318, plain, ((e2 = op(e1, e4)) | (e2 = op(e1, e3)) | (e2 = op(e1, e2)) | (e2 = op(e1, e1)) | (e2 = op(e1, e0))), inference(cnf_transformation, [], [f3])).
fof(f1224, plain, (spl10_118 | spl10_93 | spl10_68 | spl10_43 | spl10_18), inference(avatar_split_clause, [], [f319, f716, f821, f926, f1031, f1136])).
fof(f319, plain, ((e2 = op(e4, e1)) | (e2 = op(e3, e1)) | (e2 = op(e2, e1)) | (e2 = op(e1, e1)) | (e2 = op(e0, e1))), inference(cnf_transformation, [], [f3])).
fof(f1223, plain, (spl10_99 | spl10_94 | spl10_89 | spl10_84 | spl10_79), inference(avatar_split_clause, [], [f320, f972, f993, f1014, f1035, f1056])).
fof(f320, plain, ((e3 = op(e1, e4)) | (e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, e0))), inference(cnf_transformation, [], [f3])).
fof(f1222, plain, (spl10_119 | spl10_94 | spl10_69 | spl10_44 | spl10_19), inference(avatar_split_clause, [], [f321, f720, f825, f930, f1035, f1140])).
fof(f321, plain, ((e3 = op(e4, e1)) | (e3 = op(e3, e1)) | (e3 = op(e2, e1)) | (e3 = op(e1, e1)) | (e3 = op(e0, e1))), inference(cnf_transformation, [], [f3])).
fof(f1221, plain, (spl10_100 | spl10_95 | spl10_90 | spl10_85 | spl10_80), inference(avatar_split_clause, [], [f322, f976, f997, f1018, f1039, f1060])).
fof(f322, plain, ((e4 = op(e1, e4)) | (e4 = op(e1, e3)) | (e4 = op(e1, e2)) | (e4 = op(e1, e1)) | (e4 = op(e1, e0))), inference(cnf_transformation, [], [f3])).
fof(f1220, plain, (spl10_120 | spl10_95 | spl10_70 | spl10_45 | spl10_20), inference(avatar_split_clause, [], [f323, f724, f829, f934, f1039, f1144])).
fof(f323, plain, ((e4 = op(e4, e1)) | (e4 = op(e3, e1)) | (e4 = op(e2, e1)) | (e4 = op(e1, e1)) | (e4 = op(e0, e1))), inference(cnf_transformation, [], [f3])).
fof(f1219, plain, (spl10_71 | spl10_66 | spl10_61 | spl10_56 | spl10_51), inference(avatar_split_clause, [], [f324, f855, f876, f897, f918, f939])).
fof(f324, plain, ((e0 = op(e2, e4)) | (e0 = op(e2, e3)) | (e0 = op(e2, e2)) | (e0 = op(e2, e1)) | (e0 = op(e2, e0))), inference(cnf_transformation, [], [f3])).
fof(f1218, plain, (spl10_111 | spl10_86 | spl10_61 | spl10_36 | spl10_11), inference(avatar_split_clause, [], [f325, f687, f792, f897, f1002, f1107])).
fof(f325, plain, ((e0 = op(e4, e2)) | (e0 = op(e3, e2)) | (e0 = op(e2, e2)) | (e0 = op(e1, e2)) | (e0 = op(e0, e2))), inference(cnf_transformation, [], [f3])).
fof(f1217, plain, (spl10_72 | spl10_67 | spl10_62 | spl10_57 | spl10_52), inference(avatar_split_clause, [], [f326, f859, f880, f901, f922, f943])).
fof(f326, plain, ((e1 = op(e2, e4)) | (e1 = op(e2, e3)) | (e1 = op(e2, e2)) | (e1 = op(e2, e1)) | (e1 = op(e2, e0))), inference(cnf_transformation, [], [f3])).
fof(f1216, plain, (spl10_112 | spl10_87 | spl10_62 | spl10_37 | spl10_12), inference(avatar_split_clause, [], [f327, f691, f796, f901, f1006, f1111])).
fof(f327, plain, ((e1 = op(e4, e2)) | (e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(e0, e2))), inference(cnf_transformation, [], [f3])).
fof(f1213, plain, (spl10_74 | spl10_69 | spl10_64 | spl10_59 | spl10_54), inference(avatar_split_clause, [], [f330, f867, f888, f909, f930, f951])).
fof(f330, plain, ((e3 = op(e2, e4)) | (e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, e0))), inference(cnf_transformation, [], [f3])).
fof(f1212, plain, (spl10_114 | spl10_89 | spl10_64 | spl10_39 | spl10_14), inference(avatar_split_clause, [], [f331, f699, f804, f909, f1014, f1119])).
fof(f331, plain, ((e3 = op(e4, e2)) | (e3 = op(e3, e2)) | (e3 = op(e2, e2)) | (e3 = op(e1, e2)) | (e3 = op(e0, e2))), inference(cnf_transformation, [], [f3])).
fof(f1211, plain, (spl10_75 | spl10_70 | spl10_65 | spl10_60 | spl10_55), inference(avatar_split_clause, [], [f332, f871, f892, f913, f934, f955])).
fof(f332, plain, ((e4 = op(e2, e4)) | (e4 = op(e2, e3)) | (e4 = op(e2, e2)) | (e4 = op(e2, e1)) | (e4 = op(e2, e0))), inference(cnf_transformation, [], [f3])).
fof(f1210, plain, (spl10_115 | spl10_90 | spl10_65 | spl10_40 | spl10_15), inference(avatar_split_clause, [], [f333, f703, f808, f913, f1018, f1123])).
fof(f333, plain, ((e4 = op(e4, e2)) | (e4 = op(e3, e2)) | (e4 = op(e2, e2)) | (e4 = op(e1, e2)) | (e4 = op(e0, e2))), inference(cnf_transformation, [], [f3])).
fof(f1209, plain, (spl10_46 | spl10_41 | spl10_36 | spl10_31 | spl10_26), inference(avatar_split_clause, [], [f334, f750, f771, f792, f813, f834])).
fof(f334, plain, ((e0 = op(e3, e4)) | (e0 = op(e3, e3)) | (e0 = op(e3, e2)) | (e0 = op(e3, e1)) | (e0 = op(e3, e0))), inference(cnf_transformation, [], [f3])).
fof(f1208, plain, (spl10_106 | spl10_81 | spl10_56 | spl10_31 | spl10_6), inference(avatar_split_clause, [], [f335, f666, f771, f876, f981, f1086])).
fof(f335, plain, ((e0 = op(e4, e3)) | (e0 = op(e3, e3)) | (e0 = op(e2, e3)) | (e0 = op(e1, e3)) | (e0 = op(e0, e3))), inference(cnf_transformation, [], [f3])).
fof(f1207, plain, (spl10_47 | spl10_42 | spl10_37 | spl10_32 | spl10_27), inference(avatar_split_clause, [], [f336, f754, f775, f796, f817, f838])).
fof(f336, plain, ((e1 = op(e3, e4)) | (e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, e0))), inference(cnf_transformation, [], [f3])).
fof(f1206, plain, (spl10_107 | spl10_82 | spl10_57 | spl10_32 | spl10_7), inference(avatar_split_clause, [], [f337, f670, f775, f880, f985, f1090])).
fof(f337, plain, ((e1 = op(e4, e3)) | (e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(e0, e3))), inference(cnf_transformation, [], [f3])).
fof(f1205, plain, (spl10_48 | spl10_43 | spl10_38 | spl10_33 | spl10_28), inference(avatar_split_clause, [], [f338, f758, f779, f800, f821, f842])).
fof(f338, plain, ((e2 = op(e3, e4)) | (e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, e0))), inference(cnf_transformation, [], [f3])).
fof(f1204, plain, (spl10_108 | spl10_83 | spl10_58 | spl10_33 | spl10_8), inference(avatar_split_clause, [], [f339, f674, f779, f884, f989, f1094])).
fof(f339, plain, ((e2 = op(e4, e3)) | (e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(e0, e3))), inference(cnf_transformation, [], [f3])).
fof(f1201, plain, (spl10_50 | spl10_45 | spl10_40 | spl10_35 | spl10_30), inference(avatar_split_clause, [], [f342, f766, f787, f808, f829, f850])).
fof(f342, plain, ((e4 = op(e3, e4)) | (e4 = op(e3, e3)) | (e4 = op(e3, e2)) | (e4 = op(e3, e1)) | (e4 = op(e3, e0))), inference(cnf_transformation, [], [f3])).
fof(f1200, plain, (spl10_110 | spl10_85 | spl10_60 | spl10_35 | spl10_10), inference(avatar_split_clause, [], [f343, f682, f787, f892, f997, f1102])).
fof(f343, plain, ((e4 = op(e4, e3)) | (e4 = op(e3, e3)) | (e4 = op(e2, e3)) | (e4 = op(e1, e3)) | (e4 = op(e0, e3))), inference(cnf_transformation, [], [f3])).
fof(f1199, plain, (spl10_21 | spl10_16 | spl10_11 | spl10_6 | spl10_1), inference(avatar_split_clause, [], [f344, f645, f666, f687, f708, f729])).
fof(f344, plain, ((e0 = op(e4, e4)) | (e0 = op(e4, e3)) | (e0 = op(e4, e2)) | (e0 = op(e4, e1)) | (e0 = op(e4, e0))), inference(cnf_transformation, [], [f3])).
fof(f1198, plain, (spl10_101 | spl10_76 | spl10_51 | spl10_26 | spl10_1), inference(avatar_split_clause, [], [f345, f645, f750, f855, f960, f1065])).
fof(f345, plain, ((e0 = op(e4, e4)) | (e0 = op(e3, e4)) | (e0 = op(e2, e4)) | (e0 = op(e1, e4)) | (e0 = op(e0, e4))), inference(cnf_transformation, [], [f3])).
fof(f1197, plain, (spl10_22 | spl10_17 | spl10_12 | spl10_7 | spl10_2), inference(avatar_split_clause, [], [f346, f649, f670, f691, f712, f733])).
fof(f346, plain, ((e1 = op(e4, e4)) | (e1 = op(e4, e3)) | (e1 = op(e4, e2)) | (e1 = op(e4, e1)) | (e1 = op(e4, e0))), inference(cnf_transformation, [], [f3])).
fof(f1196, plain, (spl10_102 | spl10_77 | spl10_52 | spl10_27 | spl10_2), inference(avatar_split_clause, [], [f347, f649, f754, f859, f964, f1069])).
fof(f347, plain, ((e1 = op(e4, e4)) | (e1 = op(e3, e4)) | (e1 = op(e2, e4)) | (e1 = op(e1, e4)) | (e1 = op(e0, e4))), inference(cnf_transformation, [], [f3])).
fof(f1195, plain, (spl10_23 | spl10_18 | spl10_13 | spl10_8 | spl10_3), inference(avatar_split_clause, [], [f348, f653, f674, f695, f716, f737])).
fof(f348, plain, ((e2 = op(e4, e4)) | (e2 = op(e4, e3)) | (e2 = op(e4, e2)) | (e2 = op(e4, e1)) | (e2 = op(e4, e0))), inference(cnf_transformation, [], [f3])).
fof(f1194, plain, (spl10_103 | spl10_78 | spl10_53 | spl10_28 | spl10_3), inference(avatar_split_clause, [], [f349, f653, f758, f863, f968, f1073])).
fof(f349, plain, ((e2 = op(e4, e4)) | (e2 = op(e3, e4)) | (e2 = op(e2, e4)) | (e2 = op(e1, e4)) | (e2 = op(e0, e4))), inference(cnf_transformation, [], [f3])).
fof(f1193, plain, (spl10_24 | spl10_19 | spl10_14 | spl10_9 | spl10_4), inference(avatar_split_clause, [], [f350, f657, f678, f699, f720, f741])).
fof(f350, plain, ((e3 = op(e4, e4)) | (e3 = op(e4, e3)) | (e3 = op(e4, e2)) | (e3 = op(e4, e1)) | (e3 = op(e4, e0))), inference(cnf_transformation, [], [f3])).
fof(f1192, plain, (spl10_104 | spl10_79 | spl10_54 | spl10_29 | spl10_4), inference(avatar_split_clause, [], [f351, f657, f762, f867, f972, f1077])).
fof(f351, plain, ((e3 = op(e4, e4)) | (e3 = op(e3, e4)) | (e3 = op(e2, e4)) | (e3 = op(e1, e4)) | (e3 = op(e0, e4))), inference(cnf_transformation, [], [f3])).
fof(f1189, plain, (spl10_126 | spl10_127 | spl10_128 | spl10_129 | spl10_130), inference(avatar_split_clause, [], [f303, f1186, f1182, f1178, f1174, f1170])).
fof(f303, plain, ((e4 = unit) | (e3 = unit) | (e2 = unit) | (e1 = unit) | (e0 = unit)), inference(cnf_transformation, [], [f2])).
fof(f1168, plain, (spl10_121 | spl10_122 | spl10_123 | spl10_124 | spl10_125), inference(avatar_split_clause, [], [f268, f1165, f1161, f1157, f1153, f1149])).
fof(f268, plain, ((op(e0, e0) = e4) | (op(e0, e0) = e3) | (op(e0, e0) = e2) | (op(e0, e0) = e1) | (e0 = op(e0, e0))), inference(cnf_transformation, [], [f1])).
fof(f1, plain, (((e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | (e1 = op(e4, e4)) | (e0 = op(e4, e4))) & ((e4 = op(e4, e3)) | (e3 = op(e4, e3)) | (e2 = op(e4, e3)) | (e1 = op(e4, e3)) | (e0 = op(e4, e3))) & ((e4 = op(e4, e2)) | (e3 = op(e4, e2)) | (e2 = op(e4, e2)) | (e1 = op(e4, e2)) | (e0 = op(e4, e2))) & ((e4 = op(e4, e1)) | (e3 = op(e4, e1)) | (e2 = op(e4, e1)) | (e1 = op(e4, e1)) | (e0 = op(e4, e1))) & ((e4 = op(e4, e0)) | (e3 = op(e4, e0)) | (e2 = op(e4, e0)) | (e1 = op(e4, e0)) | (e0 = op(e4, e0))) & ((e4 = op(e3, e4)) | (e3 = op(e3, e4)) | (e2 = op(e3, e4)) | (e1 = op(e3, e4)) | (e0 = op(e3, e4))) & ((e4 = op(e3, e3)) | (e3 = op(e3, e3)) | (e2 = op(e3, e3)) | (e1 = op(e3, e3)) | (e0 = op(e3, e3))) & ((e4 = op(e3, e2)) | (e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (e0 = op(e3, e2))) & ((e4 = op(e3, e1)) | (e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (e0 = op(e3, e1))) & ((e4 = op(e3, e0)) | (e3 = op(e3, e0)) | (e2 = op(e3, e0)) | (e1 = op(e3, e0)) | (e0 = op(e3, e0))) & ((e4 = op(e2, e4)) | (e3 = op(e2, e4)) | (e2 = op(e2, e4)) | (e1 = op(e2, e4)) | (e0 = op(e2, e4))) & ((e4 = op(e2, e3)) | (e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))) & ((e4 = op(e2, e2)) | (e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (e0 = op(e2, e2))) & ((e4 = op(e2, e1)) | (e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))) & ((e4 = op(e2, e0)) | (e3 = op(e2, e0)) | (e2 = op(e2, e0)) | (e1 = op(e2, e0)) | (e0 = op(e2, e0))) & ((e4 = op(e1, e4)) | (e3 = op(e1, e4)) | (e2 = op(e1, e4)) | (e1 = op(e1, e4)) | (e0 = op(e1, e4))) & ((e4 = op(e1, e3)) | (e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (e0 = op(e1, e3))) & ((e4 = op(e1, e2)) | (e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (e0 = op(e1, e2))) & ((e4 = op(e1, e1)) | (e3 = op(e1, e1)) | (e2 = op(e1, e1)) | (e1 = op(e1, e1)) | (e0 = op(e1, e1))) & ((e4 = op(e1, e0)) | (e3 = op(e1, e0)) | (e2 = op(e1, e0)) | (e1 = op(e1, e0)) | (e0 = op(e1, e0))) & ((e4 = op(e0, e4)) | (e3 = op(e0, e4)) | (e2 = op(e0, e4)) | (e1 = op(e0, e4)) | (e0 = op(e0, e4))) & ((e4 = op(e0, e3)) | (e3 = op(e0, e3)) | (e2 = op(e0, e3)) | (e1 = op(e0, e3)) | (e0 = op(e0, e3))) & ((e4 = op(e0, e2)) | (e3 = op(e0, e2)) | (e2 = op(e0, e2)) | (e1 = op(e0, e2)) | (e0 = op(e0, e2))) & ((e4 = op(e0, e1)) | (e3 = op(e0, e1)) | (e2 = op(e0, e1)) | (e1 = op(e0, e1)) | (e0 = op(e0, e1))) & ((op(e0, e0) = e4) | (op(e0, e0) = e3) | (op(e0, e0) = e2) | (op(e0, e0) = e1) | (e0 = op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG098+1.p', ax1)).
fof(f1147, plain, (spl10_116 | spl10_117 | spl10_118 | spl10_119 | spl10_120), inference(avatar_split_clause, [], [f269, f1144, f1140, f1136, f1132, f1128])).
fof(f269, plain, ((e4 = op(e0, e1)) | (e3 = op(e0, e1)) | (e2 = op(e0, e1)) | (e1 = op(e0, e1)) | (e0 = op(e0, e1))), inference(cnf_transformation, [], [f1])).
fof(f1126, plain, (spl10_111 | spl10_112 | spl10_113 | spl10_114 | spl10_115), inference(avatar_split_clause, [], [f270, f1123, f1119, f1115, f1111, f1107])).
fof(f270, plain, ((e4 = op(e0, e2)) | (e3 = op(e0, e2)) | (e2 = op(e0, e2)) | (e1 = op(e0, e2)) | (e0 = op(e0, e2))), inference(cnf_transformation, [], [f1])).
fof(f1105, plain, (spl10_106 | spl10_107 | spl10_108 | spl10_109 | spl10_110), inference(avatar_split_clause, [], [f271, f1102, f1098, f1094, f1090, f1086])).
fof(f271, plain, ((e4 = op(e0, e3)) | (e3 = op(e0, e3)) | (e2 = op(e0, e3)) | (e1 = op(e0, e3)) | (e0 = op(e0, e3))), inference(cnf_transformation, [], [f1])).
fof(f1084, plain, (spl10_101 | spl10_102 | spl10_103 | spl10_104 | spl10_105), inference(avatar_split_clause, [], [f272, f1081, f1077, f1073, f1069, f1065])).
fof(f272, plain, ((e4 = op(e0, e4)) | (e3 = op(e0, e4)) | (e2 = op(e0, e4)) | (e1 = op(e0, e4)) | (e0 = op(e0, e4))), inference(cnf_transformation, [], [f1])).
fof(f1063, plain, (spl10_96 | spl10_97 | spl10_98 | spl10_99 | spl10_100), inference(avatar_split_clause, [], [f273, f1060, f1056, f1052, f1048, f1044])).
fof(f273, plain, ((e4 = op(e1, e0)) | (e3 = op(e1, e0)) | (e2 = op(e1, e0)) | (e1 = op(e1, e0)) | (e0 = op(e1, e0))), inference(cnf_transformation, [], [f1])).
fof(f1042, plain, (spl10_91 | spl10_92 | spl10_93 | spl10_94 | spl10_95), inference(avatar_split_clause, [], [f274, f1039, f1035, f1031, f1027, f1023])).
fof(f274, plain, ((e4 = op(e1, e1)) | (e3 = op(e1, e1)) | (e2 = op(e1, e1)) | (e1 = op(e1, e1)) | (e0 = op(e1, e1))), inference(cnf_transformation, [], [f1])).
fof(f1021, plain, (spl10_86 | spl10_87 | spl10_88 | spl10_89 | spl10_90), inference(avatar_split_clause, [], [f275, f1018, f1014, f1010, f1006, f1002])).
fof(f275, plain, ((e4 = op(e1, e2)) | (e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (e0 = op(e1, e2))), inference(cnf_transformation, [], [f1])).
fof(f1000, plain, (spl10_81 | spl10_82 | spl10_83 | spl10_84 | spl10_85), inference(avatar_split_clause, [], [f276, f997, f993, f989, f985, f981])).
fof(f276, plain, ((e4 = op(e1, e3)) | (e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (e0 = op(e1, e3))), inference(cnf_transformation, [], [f1])).
fof(f979, plain, (spl10_76 | spl10_77 | spl10_78 | spl10_79 | spl10_80), inference(avatar_split_clause, [], [f277, f976, f972, f968, f964, f960])).
fof(f277, plain, ((e4 = op(e1, e4)) | (e3 = op(e1, e4)) | (e2 = op(e1, e4)) | (e1 = op(e1, e4)) | (e0 = op(e1, e4))), inference(cnf_transformation, [], [f1])).
fof(f958, plain, (spl10_71 | spl10_72 | spl10_73 | spl10_74 | spl10_75), inference(avatar_split_clause, [], [f278, f955, f951, f947, f943, f939])).
fof(f278, plain, ((e4 = op(e2, e0)) | (e3 = op(e2, e0)) | (e2 = op(e2, e0)) | (e1 = op(e2, e0)) | (e0 = op(e2, e0))), inference(cnf_transformation, [], [f1])).
fof(f937, plain, (spl10_66 | spl10_67 | spl10_68 | spl10_69 | spl10_70), inference(avatar_split_clause, [], [f279, f934, f930, f926, f922, f918])).
fof(f279, plain, ((e4 = op(e2, e1)) | (e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))), inference(cnf_transformation, [], [f1])).
fof(f916, plain, (spl10_61 | spl10_62 | spl10_63 | spl10_64 | spl10_65), inference(avatar_split_clause, [], [f280, f913, f909, f905, f901, f897])).
fof(f280, plain, ((e4 = op(e2, e2)) | (e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (e0 = op(e2, e2))), inference(cnf_transformation, [], [f1])).
fof(f895, plain, (spl10_56 | spl10_57 | spl10_58 | spl10_59 | spl10_60), inference(avatar_split_clause, [], [f281, f892, f888, f884, f880, f876])).
fof(f281, plain, ((e4 = op(e2, e3)) | (e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))), inference(cnf_transformation, [], [f1])).
fof(f874, plain, (spl10_51 | spl10_52 | spl10_53 | spl10_54 | spl10_55), inference(avatar_split_clause, [], [f282, f871, f867, f863, f859, f855])).
fof(f282, plain, ((e4 = op(e2, e4)) | (e3 = op(e2, e4)) | (e2 = op(e2, e4)) | (e1 = op(e2, e4)) | (e0 = op(e2, e4))), inference(cnf_transformation, [], [f1])).
fof(f853, plain, (spl10_46 | spl10_47 | spl10_48 | spl10_49 | spl10_50), inference(avatar_split_clause, [], [f283, f850, f846, f842, f838, f834])).
fof(f283, plain, ((e4 = op(e3, e0)) | (e3 = op(e3, e0)) | (e2 = op(e3, e0)) | (e1 = op(e3, e0)) | (e0 = op(e3, e0))), inference(cnf_transformation, [], [f1])).
fof(f832, plain, (spl10_41 | spl10_42 | spl10_43 | spl10_44 | spl10_45), inference(avatar_split_clause, [], [f284, f829, f825, f821, f817, f813])).
fof(f284, plain, ((e4 = op(e3, e1)) | (e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (e0 = op(e3, e1))), inference(cnf_transformation, [], [f1])).
fof(f811, plain, (spl10_36 | spl10_37 | spl10_38 | spl10_39 | spl10_40), inference(avatar_split_clause, [], [f285, f808, f804, f800, f796, f792])).
fof(f285, plain, ((e4 = op(e3, e2)) | (e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (e0 = op(e3, e2))), inference(cnf_transformation, [], [f1])).
fof(f790, plain, (spl10_31 | spl10_32 | spl10_33 | spl10_34 | spl10_35), inference(avatar_split_clause, [], [f286, f787, f783, f779, f775, f771])).
fof(f286, plain, ((e4 = op(e3, e3)) | (e3 = op(e3, e3)) | (e2 = op(e3, e3)) | (e1 = op(e3, e3)) | (e0 = op(e3, e3))), inference(cnf_transformation, [], [f1])).
fof(f769, plain, (spl10_26 | spl10_27 | spl10_28 | spl10_29 | spl10_30), inference(avatar_split_clause, [], [f287, f766, f762, f758, f754, f750])).
fof(f287, plain, ((e4 = op(e3, e4)) | (e3 = op(e3, e4)) | (e2 = op(e3, e4)) | (e1 = op(e3, e4)) | (e0 = op(e3, e4))), inference(cnf_transformation, [], [f1])).
fof(f748, plain, (spl10_21 | spl10_22 | spl10_23 | spl10_24 | spl10_25), inference(avatar_split_clause, [], [f288, f745, f741, f737, f733, f729])).
fof(f288, plain, ((e4 = op(e4, e0)) | (e3 = op(e4, e0)) | (e2 = op(e4, e0)) | (e1 = op(e4, e0)) | (e0 = op(e4, e0))), inference(cnf_transformation, [], [f1])).
fof(f727, plain, (spl10_16 | spl10_17 | spl10_18 | spl10_19 | spl10_20), inference(avatar_split_clause, [], [f289, f724, f720, f716, f712, f708])).
fof(f289, plain, ((e4 = op(e4, e1)) | (e3 = op(e4, e1)) | (e2 = op(e4, e1)) | (e1 = op(e4, e1)) | (e0 = op(e4, e1))), inference(cnf_transformation, [], [f1])).
fof(f706, plain, (spl10_11 | spl10_12 | spl10_13 | spl10_14 | spl10_15), inference(avatar_split_clause, [], [f290, f703, f699, f695, f691, f687])).
fof(f290, plain, ((e4 = op(e4, e2)) | (e3 = op(e4, e2)) | (e2 = op(e4, e2)) | (e1 = op(e4, e2)) | (e0 = op(e4, e2))), inference(cnf_transformation, [], [f1])).
fof(f685, plain, (spl10_6 | spl10_7 | spl10_8 | spl10_9 | spl10_10), inference(avatar_split_clause, [], [f291, f682, f678, f674, f670, f666])).
fof(f291, plain, ((e4 = op(e4, e3)) | (e3 = op(e4, e3)) | (e2 = op(e4, e3)) | (e1 = op(e4, e3)) | (e0 = op(e4, e3))), inference(cnf_transformation, [], [f1])).
fof(f664, plain, (spl10_1 | spl10_2 | spl10_3 | spl10_4 | spl10_5), inference(avatar_split_clause, [], [f292, f661, f657, f653, f649, f645])).
fof(f292, plain, ((e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | (e1 = op(e4, e4)) | (e0 = op(e4, e4))), inference(cnf_transformation, [], [f1])).