fof(f20114, plain, $false, inference(avatar_sat_refutation, [], [f2339, f2514, f3539, f3697, f4671, f4674, f4677, f4680, f4683, f4686, f5138, f5165, f5192, f5216, f5240, f6396, f8165, f8283, f8285, f8286, f8287, f8289, f8290, f8292, f8703, f8775, f9214, f9469, f9748, f9752, f9874, f9883, f10074, f10124, f10151, f10153, f10294, f10323, f10335, f10427, f10470, f10510, f10529, f10580, f10583, f10639, f11245, f11973, f11978, f11995, f12045, f12124, f12528, f12724, f12770, f12810, f12863, f12864, f13157, f13160, f13288, f13396, f13509, f14197, f14577, f14837, f15062, f15362, f15496, f15880, f15978, f16223, f16395, f16424, f16551, f16572, f16661, f16663, f16705, f16706, f16876, f17003, f17042, f17122, f17150, f17229, f17321, f17347, f17394, f17648, f17836, f17998, f18093, f18169, f18216, f18257, f18302, f18614, f18837, f18983, f19049, f19118, f19121, f19218, f19276, f19328, f19381, f19426, f19471, f19518, f19566, f19613, f19664, f19719, f19769, f19819, f19870, f19912, f19957, f20005, f20056, f20104])).
fof(f20104, plain, (~ spl30_211 | ~ spl30_683 | spl30_754), inference(avatar_contradiction_clause, [], [f20103])).
fof(f20103, plain, ($false | (~ spl30_211 | ~ spl30_683 | spl30_754)), inference(subsumption_resolution, [], [f20102, f5827])).
fof(f5827, plain, ((e20 = h1(e13)) | ~ spl30_683), inference(avatar_component_clause, [], [f5826])).
fof(f5826, plain, (spl30_683 <=> (e20 = h1(e13))), introduced(avatar_definition, [new_symbols(naming, [spl30_683])])).
fof(f20102, plain, (~ (e20 = h1(e13)) | (~ spl30_211 | spl30_754)), inference(forward_demodulation, [], [f20101, f4945])).
fof(f4945, plain, (e20 = h5(e10)), inference(forward_demodulation, [], [f4944, f4915])).
fof(f4915, plain, (e20 = op2(h5(e12), h5(e13))), inference(backward_demodulation, [], [f4881, f4888])).
fof(f4888, plain, (op2(e21, e24) = h5(e12)), inference(forward_demodulation, [], [f4887, f4872])).
fof(f4872, plain, (e21 = op2(h5(e13), e24)), inference(backward_demodulation, [], [f1284, f1316])).
fof(f1316, plain, (op2(e24, e24) = h5(e13)), inference(cnf_transformation, [], [f30])).
fof(f30, plain, ((op2(op2(op2(op2(e24, e24), e24), e24), e24) = h5(e15)) & (op2(e24, e24) = h5(e13)) & (op2(op2(op2(e24, e24), e24), e24) = h5(e12)) & (op2(op2(e24, e24), e24) = h5(e11)) & (op2(op2(op2(op2(e24, e24), e24), e24), op2(e24, e24)) = h5(e10)) & (e24 = h5(e14))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG026+1.p', ax30)).
fof(f1284, plain, (e21 = op2(op2(e24, e24), e24)), inference(cnf_transformation, [], [f25])).
fof(f25, plain, ((e25 = op2(op2(op2(op2(e24, e24), e24), e24), e24)) & (e23 = op2(e24, e24)) & (e22 = op2(op2(op2(e24, e24), e24), e24)) & (e21 = op2(op2(e24, e24), e24)) & (e20 = op2(op2(op2(op2(e24, e24), e24), e24), op2(e24, e24)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG026+1.p', ax25)).
fof(f4887, plain, (h5(e12) = op2(op2(h5(e13), e24), e24)), inference(forward_demodulation, [], [f1315, f1316])).
fof(f1315, plain, (op2(op2(op2(e24, e24), e24), e24) = h5(e12)), inference(cnf_transformation, [], [f30])).
fof(f4881, plain, (e20 = op2(op2(e21, e24), h5(e13))), inference(backward_demodulation, [], [f4687, f1316])).
fof(f4687, plain, (e20 = op2(op2(e21, e24), op2(e24, e24))), inference(forward_demodulation, [], [f1283, f1284])).
fof(f1283, plain, (e20 = op2(op2(op2(op2(e24, e24), e24), e24), op2(e24, e24))), inference(cnf_transformation, [], [f25])).
fof(f4944, plain, (h5(e10) = op2(h5(e12), h5(e13))), inference(forward_demodulation, [], [f4943, f4888])).
fof(f4943, plain, (h5(e10) = op2(op2(e21, e24), h5(e13))), inference(forward_demodulation, [], [f4942, f4872])).
fof(f4942, plain, (h5(e10) = op2(op2(op2(h5(e13), e24), e24), h5(e13))), inference(forward_demodulation, [], [f1313, f1316])).
fof(f1313, plain, (op2(op2(op2(op2(e24, e24), e24), e24), op2(e24, e24)) = h5(e10)), inference(cnf_transformation, [], [f30])).
fof(f20101, plain, (~ (h1(e13) = h5(e10)) | (~ spl30_211 | spl30_754)), inference(forward_demodulation, [], [f6395, f2418])).
fof(f2418, plain, ((e10 = op1(e10, e10)) | ~ spl30_211), inference(avatar_component_clause, [], [f2416])).
fof(f2416, plain, (spl30_211 <=> (e10 = op1(e10, e10))), introduced(avatar_definition, [new_symbols(naming, [spl30_211])])).
fof(f6395, plain, (~ (h1(e13) = h5(op1(e10, e10))) | spl30_754), inference(avatar_component_clause, [], [f6393])).
fof(f6393, plain, (spl30_754 <=> (h1(e13) = h5(op1(e10, e10)))), introduced(avatar_definition, [new_symbols(naming, [spl30_754])])).
fof(f20056, plain, (~ spl30_196 | ~ spl30_454 | ~ spl30_553 | spl30_751), inference(avatar_contradiction_clause, [], [f20055])).
fof(f20055, plain, ($false | (~ spl30_196 | ~ spl30_454 | ~ spl30_553 | spl30_751)), inference(subsumption_resolution, [], [f20054, f3430])).
fof(f3430, plain, ((e23 = op2(e20, e23)) | ~ spl30_454), inference(avatar_component_clause, [], [f3428])).
fof(f3428, plain, (spl30_454 <=> (e23 = op2(e20, e23))), introduced(avatar_definition, [new_symbols(naming, [spl30_454])])).
fof(f20054, plain, (~ (e23 = op2(e20, e23)) | (~ spl30_196 | ~ spl30_553 | spl30_751)), inference(forward_demodulation, [], [f20053, f5163])).
fof(f5163, plain, ((e23 = h5(e13)) | ~ spl30_553), inference(avatar_component_clause, [], [f5162])).
fof(f5162, plain, (spl30_553 <=> (e23 = h5(e13))), introduced(avatar_definition, [new_symbols(naming, [spl30_553])])).
fof(f20053, plain, (~ (op2(e20, e23) = h5(e13)) | (~ spl30_196 | ~ spl30_553 | spl30_751)), inference(forward_demodulation, [], [f20052, f2355])).
fof(f2355, plain, ((e13 = op1(e10, e13)) | ~ spl30_196), inference(avatar_component_clause, [], [f2353])).
fof(f2353, plain, (spl30_196 <=> (e13 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl30_196])])).
fof(f20052, plain, (~ (op2(e20, e23) = h5(op1(e10, e13))) | (~ spl30_553 | spl30_751)), inference(forward_demodulation, [], [f6383, f5163])).
fof(f6383, plain, (~ (h5(op1(e10, e13)) = op2(e20, h5(e13))) | spl30_751), inference(avatar_component_clause, [], [f6381])).
fof(f6381, plain, (spl30_751 <=> (h5(op1(e10, e13)) = op2(e20, h5(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl30_751])])).
fof(f20005, plain, (~ spl30_186 | ~ spl30_444 | ~ spl30_732 | spl30_749), inference(avatar_contradiction_clause, [], [f20004])).
fof(f20004, plain, ($false | (~ spl30_186 | ~ spl30_444 | ~ spl30_732 | spl30_749)), inference(subsumption_resolution, [], [f20003, f3388])).
fof(f3388, plain, ((e25 = op2(e20, e25)) | ~ spl30_444), inference(avatar_component_clause, [], [f3386])).
fof(f3386, plain, (spl30_444 <=> (e25 = op2(e20, e25))), introduced(avatar_definition, [new_symbols(naming, [spl30_444])])).
fof(f20003, plain, (~ (e25 = op2(e20, e25)) | (~ spl30_186 | ~ spl30_732 | spl30_749)), inference(forward_demodulation, [], [f20002, f6306])).
fof(f6306, plain, ((e25 = h5(e15)) | ~ spl30_732), inference(avatar_component_clause, [], [f6305])).
fof(f6305, plain, (spl30_732 <=> (e25 = h5(e15))), introduced(avatar_definition, [new_symbols(naming, [spl30_732])])).
fof(f20002, plain, (~ (op2(e20, e25) = h5(e15)) | (~ spl30_186 | ~ spl30_732 | spl30_749)), inference(forward_demodulation, [], [f20001, f2313])).
fof(f2313, plain, ((e15 = op1(e10, e15)) | ~ spl30_186), inference(avatar_component_clause, [], [f2311])).
fof(f2311, plain, (spl30_186 <=> (e15 = op1(e10, e15))), introduced(avatar_definition, [new_symbols(naming, [spl30_186])])).
fof(f20001, plain, (~ (op2(e20, e25) = h5(op1(e10, e15))) | (~ spl30_732 | spl30_749)), inference(forward_demodulation, [], [f6375, f6306])).
fof(f6375, plain, (~ (h5(op1(e10, e15)) = op2(e20, h5(e15))) | spl30_749), inference(avatar_component_clause, [], [f6373])).
fof(f6373, plain, (spl30_749 <=> (h5(op1(e10, e15)) = op2(e20, h5(e15)))), introduced(avatar_definition, [new_symbols(naming, [spl30_749])])).
fof(f19957, plain, (~ spl30_167 | ~ spl30_425 | ~ spl30_558 | spl30_747), inference(avatar_contradiction_clause, [], [f19956])).
fof(f19956, plain, ($false | (~ spl30_167 | ~ spl30_425 | ~ spl30_558 | spl30_747)), inference(subsumption_resolution, [], [f19955, f3309])).
fof(f3309, plain, ((e24 = op2(e21, e22)) | ~ spl30_425), inference(avatar_component_clause, [], [f3307])).
fof(f3307, plain, (spl30_425 <=> (e24 = op2(e21, e22))), introduced(avatar_definition, [new_symbols(naming, [spl30_425])])).
fof(f19955, plain, (~ (e24 = op2(e21, e22)) | (~ spl30_167 | ~ spl30_558 | spl30_747)), inference(forward_demodulation, [], [f19954, f1312])).
fof(f1312, plain, (e24 = h5(e14)), inference(cnf_transformation, [], [f30])).
fof(f19954, plain, (~ (op2(e21, e22) = h5(e14)) | (~ spl30_167 | ~ spl30_558 | spl30_747)), inference(forward_demodulation, [], [f19953, f2234])).
fof(f2234, plain, ((e14 = op1(e11, e12)) | ~ spl30_167), inference(avatar_component_clause, [], [f2232])).
fof(f2232, plain, (spl30_167 <=> (e14 = op1(e11, e12))), introduced(avatar_definition, [new_symbols(naming, [spl30_167])])).
fof(f19953, plain, (~ (op2(e21, e22) = h5(op1(e11, e12))) | (~ spl30_558 | spl30_747)), inference(forward_demodulation, [], [f6367, f5190])).
fof(f5190, plain, ((e22 = h5(e12)) | ~ spl30_558), inference(avatar_component_clause, [], [f5189])).
fof(f5189, plain, (spl30_558 <=> (e22 = h5(e12))), introduced(avatar_definition, [new_symbols(naming, [spl30_558])])).
fof(f6367, plain, (~ (h5(op1(e11, e12)) = op2(e21, h5(e12))) | spl30_747), inference(avatar_component_clause, [], [f6365])).
fof(f6365, plain, (spl30_747 <=> (h5(op1(e11, e12)) = op2(e21, h5(e12)))), introduced(avatar_definition, [new_symbols(naming, [spl30_747])])).
fof(f19912, plain, (~ spl30_121 | spl30_743), inference(avatar_contradiction_clause, [], [f19911])).
fof(f19911, plain, ($false | (~ spl30_121 | spl30_743)), inference(subsumption_resolution, [], [f19910, f4945])).
fof(f19910, plain, (~ (e20 = h5(e10)) | (~ spl30_121 | spl30_743)), inference(forward_demodulation, [], [f6351, f2043])).
fof(f2043, plain, ((e10 = op1(e12, e13)) | ~ spl30_121), inference(avatar_component_clause, [], [f2041])).
fof(f2041, plain, (spl30_121 <=> (e10 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl30_121])])).
fof(f6351, plain, (~ (e20 = h5(op1(e12, e13))) | spl30_743), inference(avatar_component_clause, [], [f6349])).
fof(f6349, plain, (spl30_743 <=> (e20 = h5(op1(e12, e13)))), introduced(avatar_definition, [new_symbols(naming, [spl30_743])])).
fof(f19870, plain, (~ spl30_121 | ~ spl30_349 | ~ spl30_553 | ~ spl30_558 | spl30_741), inference(avatar_contradiction_clause, [], [f19869])).
fof(f19869, plain, ($false | (~ spl30_121 | ~ spl30_349 | ~ spl30_553 | ~ spl30_558 | spl30_741)), inference(subsumption_resolution, [], [f19868, f2993])).
fof(f2993, plain, ((e20 = op2(e23, e22)) | ~ spl30_349), inference(avatar_component_clause, [], [f2991])).
fof(f2991, plain, (spl30_349 <=> (e20 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl30_349])])).
fof(f19868, plain, (~ (e20 = op2(e23, e22)) | (~ spl30_121 | ~ spl30_553 | ~ spl30_558 | spl30_741)), inference(forward_demodulation, [], [f19867, f4945])).
fof(f19867, plain, (~ (op2(e23, e22) = h5(e10)) | (~ spl30_121 | ~ spl30_553 | ~ spl30_558 | spl30_741)), inference(forward_demodulation, [], [f19866, f2043])).
fof(f19866, plain, (~ (op2(e23, e22) = h5(op1(e12, e13))) | (~ spl30_553 | ~ spl30_558 | spl30_741)), inference(forward_demodulation, [], [f19865, f5163])).
fof(f19865, plain, (~ (h5(op1(e12, e13)) = op2(h5(e13), e22)) | (~ spl30_558 | spl30_741)), inference(forward_demodulation, [], [f6343, f5190])).
fof(f6343, plain, (~ (h5(op1(e12, e13)) = op2(h5(e13), h5(e12))) | spl30_741), inference(avatar_component_clause, [], [f6341])).
fof(f6341, plain, (spl30_741 <=> (h5(op1(e12, e13)) = op2(h5(e13), h5(e12)))), introduced(avatar_definition, [new_symbols(naming, [spl30_741])])).
fof(f19819, plain, (~ spl30_46 | spl30_739), inference(avatar_contradiction_clause, [], [f19818])).
fof(f19818, plain, ($false | (~ spl30_46 | spl30_739)), inference(trivial_inequality_removal, [], [f19817])).
fof(f19817, plain, (~ (h5(e13) = h5(e13)) | (~ spl30_46 | spl30_739)), inference(forward_demodulation, [], [f6335, f1730])).
fof(f1730, plain, ((e13 = op1(e14, e14)) | ~ spl30_46), inference(avatar_component_clause, [], [f1728])).
fof(f1728, plain, (spl30_46 <=> (e13 = op1(e14, e14))), introduced(avatar_definition, [new_symbols(naming, [spl30_46])])).
fof(f6335, plain, (~ (h5(e13) = h5(op1(e14, e14))) | spl30_739), inference(avatar_component_clause, [], [f6333])).
fof(f6333, plain, (spl30_739 <=> (h5(e13) = h5(op1(e14, e14)))), introduced(avatar_definition, [new_symbols(naming, [spl30_739])])).
fof(f19769, plain, (~ spl30_186 | ~ spl30_294 | ~ spl30_732 | spl30_737), inference(avatar_contradiction_clause, [], [f19768])).
fof(f19768, plain, ($false | (~ spl30_186 | ~ spl30_294 | ~ spl30_732 | spl30_737)), inference(subsumption_resolution, [], [f19767, f2763])).
fof(f2763, plain, ((e25 = op2(e25, e20)) | ~ spl30_294), inference(avatar_component_clause, [], [f2761])).
fof(f2761, plain, (spl30_294 <=> (e25 = op2(e25, e20))), introduced(avatar_definition, [new_symbols(naming, [spl30_294])])).
fof(f19767, plain, (~ (e25 = op2(e25, e20)) | (~ spl30_186 | ~ spl30_732 | spl30_737)), inference(forward_demodulation, [], [f19766, f6306])).
fof(f19766, plain, (~ (op2(e25, e20) = h5(e15)) | (~ spl30_186 | ~ spl30_732 | spl30_737)), inference(forward_demodulation, [], [f19765, f2313])).
fof(f19765, plain, (~ (op2(e25, e20) = h5(op1(e10, e15))) | (~ spl30_732 | spl30_737)), inference(forward_demodulation, [], [f6327, f6306])).
fof(f6327, plain, (~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | spl30_737), inference(avatar_component_clause, [], [f6325])).
fof(f6325, plain, (spl30_737 <=> (h5(op1(e10, e15)) = op2(h5(e15), e20))), introduced(avatar_definition, [new_symbols(naming, [spl30_737])])).
fof(f19719, plain, (~ spl30_110 | ~ spl30_278 | ~ spl30_558 | ~ spl30_732 | spl30_735), inference(avatar_contradiction_clause, [], [f19718])).
fof(f19718, plain, ($false | (~ spl30_110 | ~ spl30_278 | ~ spl30_558 | ~ spl30_732 | spl30_735)), inference(subsumption_resolution, [], [f19717, f2697])).
fof(f2697, plain, ((e21 = op2(e25, e22)) | ~ spl30_278), inference(avatar_component_clause, [], [f2695])).
fof(f2695, plain, (spl30_278 <=> (e21 = op2(e25, e22))), introduced(avatar_definition, [new_symbols(naming, [spl30_278])])).
fof(f19717, plain, (~ (e21 = op2(e25, e22)) | (~ spl30_110 | ~ spl30_558 | ~ spl30_732 | spl30_735)), inference(forward_demodulation, [], [f19716, f4941])).
fof(f4941, plain, (e21 = h5(e11)), inference(forward_demodulation, [], [f4940, f4872])).
fof(f4940, plain, (h5(e11) = op2(h5(e13), e24)), inference(forward_demodulation, [], [f1314, f1316])).
fof(f1314, plain, (op2(op2(e24, e24), e24) = h5(e11)), inference(cnf_transformation, [], [f30])).
fof(f19716, plain, (~ (op2(e25, e22) = h5(e11)) | (~ spl30_110 | ~ spl30_558 | ~ spl30_732 | spl30_735)), inference(forward_demodulation, [], [f19715, f1997])).
fof(f1997, plain, ((e11 = op1(e12, e15)) | ~ spl30_110), inference(avatar_component_clause, [], [f1995])).
fof(f1995, plain, (spl30_110 <=> (e11 = op1(e12, e15))), introduced(avatar_definition, [new_symbols(naming, [spl30_110])])).
fof(f19715, plain, (~ (op2(e25, e22) = h5(op1(e12, e15))) | (~ spl30_558 | ~ spl30_732 | spl30_735)), inference(forward_demodulation, [], [f19714, f6306])).
fof(f19714, plain, (~ (h5(op1(e12, e15)) = op2(h5(e15), e22)) | (~ spl30_558 | spl30_735)), inference(forward_demodulation, [], [f6319, f5190])).
fof(f6319, plain, (~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | spl30_735), inference(avatar_component_clause, [], [f6317])).
fof(f6317, plain, (spl30_735 <=> (h5(op1(e12, e15)) = op2(h5(e15), h5(e12)))), introduced(avatar_definition, [new_symbols(naming, [spl30_735])])).
fof(f19664, plain, (~ spl30_206 | ~ spl30_464 | spl30_753), inference(avatar_contradiction_clause, [], [f19663])).
fof(f19663, plain, ($false | (~ spl30_206 | ~ spl30_464 | spl30_753)), inference(subsumption_resolution, [], [f19662, f3472])).
fof(f3472, plain, ((e21 = op2(e20, e21)) | ~ spl30_464), inference(avatar_component_clause, [], [f3470])).
fof(f3470, plain, (spl30_464 <=> (e21 = op2(e20, e21))), introduced(avatar_definition, [new_symbols(naming, [spl30_464])])).
fof(f19662, plain, (~ (e21 = op2(e20, e21)) | (~ spl30_206 | spl30_753)), inference(forward_demodulation, [], [f19661, f4941])).
fof(f19661, plain, (~ (op2(e20, e21) = h5(e11)) | (~ spl30_206 | spl30_753)), inference(forward_demodulation, [], [f6391, f2397])).
fof(f2397, plain, ((e11 = op1(e10, e11)) | ~ spl30_206), inference(avatar_component_clause, [], [f2395])).
fof(f2395, plain, (spl30_206 <=> (e11 = op1(e10, e11))), introduced(avatar_definition, [new_symbols(naming, [spl30_206])])).
fof(f6391, plain, (~ (op2(e20, e21) = h5(op1(e10, e11))) | spl30_753), inference(avatar_component_clause, [], [f6389])).
fof(f6389, plain, (spl30_753 <=> (op2(e20, e21) = h5(op1(e10, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl30_753])])).
fof(f19613, plain, (~ spl30_191 | ~ spl30_449 | spl30_750), inference(avatar_contradiction_clause, [], [f19612])).
fof(f19612, plain, ($false | (~ spl30_191 | ~ spl30_449 | spl30_750)), inference(subsumption_resolution, [], [f19611, f3409])).
fof(f3409, plain, ((e24 = op2(e20, e24)) | ~ spl30_449), inference(avatar_component_clause, [], [f3407])).
fof(f3407, plain, (spl30_449 <=> (e24 = op2(e20, e24))), introduced(avatar_definition, [new_symbols(naming, [spl30_449])])).
fof(f19611, plain, (~ (e24 = op2(e20, e24)) | (~ spl30_191 | spl30_750)), inference(forward_demodulation, [], [f19610, f1312])).
fof(f19610, plain, (~ (op2(e20, e24) = h5(e14)) | (~ spl30_191 | spl30_750)), inference(forward_demodulation, [], [f6379, f2334])).
fof(f2334, plain, ((e14 = op1(e10, e14)) | ~ spl30_191), inference(avatar_component_clause, [], [f2332])).
fof(f2332, plain, (spl30_191 <=> (e14 = op1(e10, e14))), introduced(avatar_definition, [new_symbols(naming, [spl30_191])])).
fof(f6379, plain, (~ (op2(e20, e24) = h5(op1(e10, e14))) | spl30_750), inference(avatar_component_clause, [], [f6377])).
fof(f6377, plain, (spl30_750 <=> (op2(e20, e24) = h5(op1(e10, e14)))), introduced(avatar_definition, [new_symbols(naming, [spl30_750])])).
fof(f19566, plain, (~ spl30_169 | ~ spl30_653 | spl30_748), inference(avatar_contradiction_clause, [], [f19565])).
fof(f19565, plain, ($false | (~ spl30_169 | ~ spl30_653 | spl30_748)), inference(subsumption_resolution, [], [f19564, f5676])).
fof(f5676, plain, ((e20 = h2(e13)) | ~ spl30_653), inference(avatar_component_clause, [], [f5675])).
fof(f5675, plain, (spl30_653 <=> (e20 = h2(e13))), introduced(avatar_definition, [new_symbols(naming, [spl30_653])])).
fof(f19564, plain, (~ (e20 = h2(e13)) | (~ spl30_169 | spl30_748)), inference(forward_demodulation, [], [f19563, f4945])).
fof(f19563, plain, (~ (h2(e13) = h5(e10)) | (~ spl30_169 | spl30_748)), inference(forward_demodulation, [], [f6371, f2243])).
fof(f2243, plain, ((e10 = op1(e11, e11)) | ~ spl30_169), inference(avatar_component_clause, [], [f2241])).
fof(f2241, plain, (spl30_169 <=> (e10 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl30_169])])).
fof(f6371, plain, (~ (h2(e13) = h5(op1(e11, e11))) | spl30_748), inference(avatar_component_clause, [], [f6369])).
fof(f6369, plain, (spl30_748 <=> (h2(e13) = h5(op1(e11, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl30_748])])).
fof(f19518, plain, (~ spl30_162 | spl30_746), inference(avatar_contradiction_clause, [], [f19517])).
fof(f19517, plain, ($false | (~ spl30_162 | spl30_746)), inference(trivial_inequality_removal, [], [f19516])).
fof(f19516, plain, (~ (h5(e15) = h5(e15)) | (~ spl30_162 | spl30_746)), inference(forward_demodulation, [], [f6363, f2213])).
fof(f2213, plain, ((e15 = op1(e11, e13)) | ~ spl30_162), inference(avatar_component_clause, [], [f2211])).
fof(f2211, plain, (spl30_162 <=> (e15 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl30_162])])).
fof(f6363, plain, (~ (h5(e15) = h5(op1(e11, e13))) | spl30_746), inference(avatar_component_clause, [], [f6361])).
fof(f6361, plain, (spl30_746 <=> (h5(e15) = h5(op1(e11, e13)))), introduced(avatar_definition, [new_symbols(naming, [spl30_746])])).
fof(f19471, plain, (~ spl30_120 | spl30_742), inference(avatar_contradiction_clause, [], [f19470])).
fof(f19470, plain, ($false | (~ spl30_120 | spl30_742)), inference(trivial_inequality_removal, [], [f19469])).
fof(f19469, plain, (~ (h5(e15) = h5(e15)) | (~ spl30_120 | spl30_742)), inference(forward_demodulation, [], [f6347, f2038])).
fof(f2038, plain, ((e15 = op1(e12, e14)) | ~ spl30_120), inference(avatar_component_clause, [], [f2036])).
fof(f2036, plain, (spl30_120 <=> (e15 = op1(e12, e14))), introduced(avatar_definition, [new_symbols(naming, [spl30_120])])).
fof(f6347, plain, (~ (h5(e15) = h5(op1(e12, e14))) | spl30_742), inference(avatar_component_clause, [], [f6345])).
fof(f6345, plain, (spl30_742 <=> (h5(e15) = h5(op1(e12, e14)))), introduced(avatar_definition, [new_symbols(naming, [spl30_742])])).
fof(f19426, plain, (~ spl30_80 | spl30_740), inference(avatar_contradiction_clause, [], [f19425])).
fof(f19425, plain, ($false | (~ spl30_80 | spl30_740)), inference(subsumption_resolution, [], [f19424, f4941])).
fof(f19424, plain, (~ (e21 = h5(e11)) | (~ spl30_80 | spl30_740)), inference(forward_demodulation, [], [f6339, f1872])).
fof(f1872, plain, ((e11 = op1(e13, e14)) | ~ spl30_80), inference(avatar_component_clause, [], [f1870])).
fof(f1870, plain, (spl30_80 <=> (e11 = op1(e13, e14))), introduced(avatar_definition, [new_symbols(naming, [spl30_80])])).
fof(f6339, plain, (~ (e21 = h5(op1(e13, e14))) | spl30_740), inference(avatar_component_clause, [], [f6337])).
fof(f6337, plain, (spl30_740 <=> (e21 = h5(op1(e13, e14)))), introduced(avatar_definition, [new_symbols(naming, [spl30_740])])).
fof(f19381, plain, (~ spl30_37 | ~ spl30_295 | ~ spl30_732 | spl30_738), inference(avatar_contradiction_clause, [], [f19380])).
fof(f19380, plain, ($false | (~ spl30_37 | ~ spl30_295 | ~ spl30_732 | spl30_738)), inference(subsumption_resolution, [], [f19379, f2768])).
fof(f2768, plain, ((e20 = op2(e24, e25)) | ~ spl30_295), inference(avatar_component_clause, [], [f2766])).
fof(f2766, plain, (spl30_295 <=> (e20 = op2(e24, e25))), introduced(avatar_definition, [new_symbols(naming, [spl30_295])])).
fof(f19379, plain, (~ (e20 = op2(e24, e25)) | (~ spl30_37 | ~ spl30_732 | spl30_738)), inference(forward_demodulation, [], [f19378, f4945])).
fof(f19378, plain, (~ (op2(e24, e25) = h5(e10)) | (~ spl30_37 | ~ spl30_732 | spl30_738)), inference(forward_demodulation, [], [f19377, f1693])).
fof(f1693, plain, ((e10 = op1(e14, e15)) | ~ spl30_37), inference(avatar_component_clause, [], [f1691])).
fof(f1691, plain, (spl30_37 <=> (e10 = op1(e14, e15))), introduced(avatar_definition, [new_symbols(naming, [spl30_37])])).
fof(f19377, plain, (~ (op2(e24, e25) = h5(op1(e14, e15))) | (~ spl30_732 | spl30_738)), inference(forward_demodulation, [], [f6331, f6306])).
fof(f6331, plain, (~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | spl30_738), inference(avatar_component_clause, [], [f6329])).
fof(f6329, plain, (spl30_738 <=> (h5(op1(e14, e15)) = op2(e24, h5(e15)))), introduced(avatar_definition, [new_symbols(naming, [spl30_738])])).
fof(f19328, plain, (~ spl30_148 | ~ spl30_286 | ~ spl30_553 | ~ spl30_732 | spl30_736), inference(avatar_contradiction_clause, [], [f19327])).
fof(f19327, plain, ($false | (~ spl30_148 | ~ spl30_286 | ~ spl30_553 | ~ spl30_732 | spl30_736)), inference(subsumption_resolution, [], [f19326, f2730])).
fof(f2730, plain, ((e23 = op2(e25, e21)) | ~ spl30_286), inference(avatar_component_clause, [], [f2728])).
fof(f2728, plain, (spl30_286 <=> (e23 = op2(e25, e21))), introduced(avatar_definition, [new_symbols(naming, [spl30_286])])).
fof(f19326, plain, (~ (e23 = op2(e25, e21)) | (~ spl30_148 | ~ spl30_553 | ~ spl30_732 | spl30_736)), inference(forward_demodulation, [], [f19325, f5163])).
fof(f19325, plain, (~ (op2(e25, e21) = h5(e13)) | (~ spl30_148 | ~ spl30_732 | spl30_736)), inference(forward_demodulation, [], [f19324, f2155])).
fof(f2155, plain, ((e13 = op1(e11, e15)) | ~ spl30_148), inference(avatar_component_clause, [], [f2153])).
fof(f2153, plain, (spl30_148 <=> (e13 = op1(e11, e15))), introduced(avatar_definition, [new_symbols(naming, [spl30_148])])).
fof(f19324, plain, (~ (op2(e25, e21) = h5(op1(e11, e15))) | (~ spl30_732 | spl30_736)), inference(forward_demodulation, [], [f6323, f6306])).
fof(f6323, plain, (~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | spl30_736), inference(avatar_component_clause, [], [f6321])).
fof(f6321, plain, (spl30_736 <=> (h5(op1(e11, e15)) = op2(h5(e15), e21))), introduced(avatar_definition, [new_symbols(naming, [spl30_736])])).
fof(f19276, plain, (~ spl30_77 | ~ spl30_275 | ~ spl30_553 | ~ spl30_732 | spl30_734), inference(avatar_contradiction_clause, [], [f19275])).
fof(f19275, plain, ($false | (~ spl30_77 | ~ spl30_275 | ~ spl30_553 | ~ spl30_732 | spl30_734)), inference(subsumption_resolution, [], [f19274, f2684])).
fof(f2684, plain, ((e24 = op2(e25, e23)) | ~ spl30_275), inference(avatar_component_clause, [], [f2682])).
fof(f2682, plain, (spl30_275 <=> (e24 = op2(e25, e23))), introduced(avatar_definition, [new_symbols(naming, [spl30_275])])).
fof(f19274, plain, (~ (e24 = op2(e25, e23)) | (~ spl30_77 | ~ spl30_553 | ~ spl30_732 | spl30_734)), inference(forward_demodulation, [], [f19273, f1312])).
fof(f19273, plain, (~ (op2(e25, e23) = h5(e14)) | (~ spl30_77 | ~ spl30_553 | ~ spl30_732 | spl30_734)), inference(forward_demodulation, [], [f19272, f1859])).
fof(f1859, plain, ((e14 = op1(e13, e15)) | ~ spl30_77), inference(avatar_component_clause, [], [f1857])).
fof(f1857, plain, (spl30_77 <=> (e14 = op1(e13, e15))), introduced(avatar_definition, [new_symbols(naming, [spl30_77])])).
fof(f19272, plain, (~ (op2(e25, e23) = h5(op1(e13, e15))) | (~ spl30_553 | ~ spl30_732 | spl30_734)), inference(forward_demodulation, [], [f19271, f6306])).
fof(f19271, plain, (~ (h5(op1(e13, e15)) = op2(h5(e15), e23)) | (~ spl30_553 | spl30_734)), inference(forward_demodulation, [], [f6315, f5163])).
fof(f6315, plain, (~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | spl30_734), inference(avatar_component_clause, [], [f6313])).
fof(f6313, plain, (spl30_734 <=> (h5(op1(e13, e15)) = op2(h5(e15), h5(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl30_734])])).
fof(f19218, plain, (~ spl30_37 | ~ spl30_265 | ~ spl30_732 | spl30_733), inference(avatar_contradiction_clause, [], [f19217])).
fof(f19217, plain, ($false | (~ spl30_37 | ~ spl30_265 | ~ spl30_732 | spl30_733)), inference(subsumption_resolution, [], [f19216, f2643])).
fof(f2643, plain, ((e20 = op2(e25, e24)) | ~ spl30_265), inference(avatar_component_clause, [], [f2641])).
fof(f2641, plain, (spl30_265 <=> (e20 = op2(e25, e24))), introduced(avatar_definition, [new_symbols(naming, [spl30_265])])).
fof(f19216, plain, (~ (e20 = op2(e25, e24)) | (~ spl30_37 | ~ spl30_732 | spl30_733)), inference(forward_demodulation, [], [f19215, f4945])).
fof(f19215, plain, (~ (op2(e25, e24) = h5(e10)) | (~ spl30_37 | ~ spl30_732 | spl30_733)), inference(forward_demodulation, [], [f19214, f1693])).
fof(f19214, plain, (~ (op2(e25, e24) = h5(op1(e14, e15))) | (~ spl30_732 | spl30_733)), inference(forward_demodulation, [], [f6311, f6306])).
fof(f6311, plain, (~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | spl30_733), inference(avatar_component_clause, [], [f6309])).
fof(f6309, plain, (spl30_733 <=> (h5(op1(e14, e15)) = op2(h5(e15), e24))), introduced(avatar_definition, [new_symbols(naming, [spl30_733])])).
fof(f19121, plain, (spl30_261 | ~ spl30_295 | ~ spl30_378 | ~ spl30_399), inference(avatar_split_clause, [], [f19120, f3199, f3111, f2766, f2624])).
fof(f2624, plain, (spl30_261 <=> (e22 = op2(e25, e25))), introduced(avatar_definition, [new_symbols(naming, [spl30_261])])).
fof(f3111, plain, (spl30_378 <=> (e25 = op2(e22, e24))), introduced(avatar_definition, [new_symbols(naming, [spl30_378])])).
fof(f3199, plain, (spl30_399 <=> (e22 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl30_399])])).
fof(f19120, plain, ((e22 = op2(e25, e25)) | (~ spl30_295 | ~ spl30_378 | ~ spl30_399)), inference(forward_demodulation, [], [f19119, f3201])).
fof(f3201, plain, ((e22 = op2(e22, e20)) | ~ spl30_399), inference(avatar_component_clause, [], [f3199])).
fof(f19119, plain, ((op2(e22, e20) = op2(e25, e25)) | (~ spl30_295 | ~ spl30_378)), inference(forward_demodulation, [], [f8792, f2768])).
fof(f8792, plain, ((op2(e25, e25) = op2(e22, op2(e24, e25))) | ~ spl30_378), inference(backward_demodulation, [], [f518, f3113])).
fof(f3113, plain, ((e25 = op2(e22, e24)) | ~ spl30_378), inference(avatar_component_clause, [], [f3111])).
fof(f518, plain, (op2(op2(e22, e24), e25) = op2(e22, op2(e24, e25))), inference(cnf_transformation, [], [f6])).
fof(f6, plain, ((op2(op2(e25, e25), e25) = op2(e25, op2(e25, e25))) & (op2(op2(e25, e25), e24) = op2(e25, op2(e25, e24))) & (op2(op2(e25, e25), e23) = op2(e25, op2(e25, e23))) & (op2(op2(e25, e25), e22) = op2(e25, op2(e25, e22))) & (op2(op2(e25, e25), e21) = op2(e25, op2(e25, e21))) & (op2(op2(e25, e25), e20) = op2(e25, op2(e25, e20))) & (op2(op2(e25, e24), e25) = op2(e25, op2(e24, e25))) & (op2(op2(e25, e24), e24) = op2(e25, op2(e24, e24))) & (op2(op2(e25, e24), e23) = op2(e25, op2(e24, e23))) & (op2(op2(e25, e24), e22) = op2(e25, op2(e24, e22))) & (op2(op2(e25, e24), e21) = op2(e25, op2(e24, e21))) & (op2(op2(e25, e24), e20) = op2(e25, op2(e24, e20))) & (op2(op2(e25, e23), e25) = op2(e25, op2(e23, e25))) & (op2(op2(e25, e23), e24) = op2(e25, op2(e23, e24))) & (op2(op2(e25, e23), e23) = op2(e25, op2(e23, e23))) & (op2(op2(e25, e23), e22) = op2(e25, op2(e23, e22))) & (op2(op2(e25, e23), e21) = op2(e25, op2(e23, e21))) & (op2(op2(e25, e23), e20) = op2(e25, op2(e23, e20))) & (op2(op2(e25, e22), e25) = op2(e25, op2(e22, e25))) & (op2(op2(e25, e22), e24) = op2(e25, op2(e22, e24))) & (op2(op2(e25, e22), e23) = op2(e25, op2(e22, e23))) & (op2(op2(e25, e22), e22) = op2(e25, op2(e22, e22))) & (op2(op2(e25, e22), e21) = op2(e25, op2(e22, e21))) & (op2(op2(e25, e22), e20) = op2(e25, op2(e22, e20))) & (op2(op2(e25, e21), e25) = op2(e25, op2(e21, e25))) & (op2(op2(e25, e21), e24) = op2(e25, op2(e21, e24))) & (op2(op2(e25, e21), e23) = op2(e25, op2(e21, e23))) & (op2(op2(e25, e21), e22) = op2(e25, op2(e21, e22))) & (op2(op2(e25, e21), e21) = op2(e25, op2(e21, e21))) & (op2(op2(e25, e21), e20) = op2(e25, op2(e21, e20))) & (op2(op2(e25, e20), e25) = op2(e25, op2(e20, e25))) & (op2(op2(e25, e20), e24) = op2(e25, op2(e20, e24))) & (op2(op2(e25, e20), e23) = op2(e25, op2(e20, e23))) & (op2(op2(e25, e20), e22) = op2(e25, op2(e20, e22))) & (op2(op2(e25, e20), e21) = op2(e25, op2(e20, e21))) & (op2(op2(e25, e20), e20) = op2(e25, op2(e20, e20))) & (op2(op2(e24, e25), e25) = op2(e24, op2(e25, e25))) & (op2(op2(e24, e25), e24) = op2(e24, op2(e25, e24))) & (op2(op2(e24, e25), e23) = op2(e24, op2(e25, e23))) & (op2(op2(e24, e25), e22) = op2(e24, op2(e25, e22))) & (op2(op2(e24, e25), e21) = op2(e24, op2(e25, e21))) & (op2(op2(e24, e25), e20) = op2(e24, op2(e25, e20))) & (op2(op2(e24, e24), e25) = op2(e24, op2(e24, e25))) & (op2(op2(e24, e24), e24) = op2(e24, op2(e24, e24))) & (op2(op2(e24, e24), e23) = op2(e24, op2(e24, e23))) & (op2(op2(e24, e24), e22) = op2(e24, op2(e24, e22))) & (op2(op2(e24, e24), e21) = op2(e24, op2(e24, e21))) & (op2(op2(e24, e24), e20) = op2(e24, op2(e24, e20))) & (op2(op2(e24, e23), e25) = op2(e24, op2(e23, e25))) & (op2(op2(e24, e23), e24) = op2(e24, op2(e23, e24))) & (op2(op2(e24, e23), e23) = op2(e24, op2(e23, e23))) & (op2(op2(e24, e23), e22) = op2(e24, op2(e23, e22))) & (op2(op2(e24, e23), e21) = op2(e24, op2(e23, e21))) & (op2(op2(e24, e23), e20) = op2(e24, op2(e23, e20))) & (op2(op2(e24, e22), e25) = op2(e24, op2(e22, e25))) & (op2(op2(e24, e22), e24) = op2(e24, op2(e22, e24))) & (op2(op2(e24, e22), e23) = op2(e24, op2(e22, e23))) & (op2(op2(e24, e22), e22) = op2(e24, op2(e22, e22))) & (op2(op2(e24, e22), e21) = op2(e24, op2(e22, e21))) & (op2(op2(e24, e22), e20) = op2(e24, op2(e22, e20))) & (op2(op2(e24, e21), e25) = op2(e24, op2(e21, e25))) & (op2(op2(e24, e21), e24) = op2(e24, op2(e21, e24))) & (op2(op2(e24, e21), e23) = op2(e24, op2(e21, e23))) & (op2(op2(e24, e21), e22) = op2(e24, op2(e21, e22))) & (op2(op2(e24, e21), e21) = op2(e24, op2(e21, e21))) & (op2(op2(e24, e21), e20) = op2(e24, op2(e21, e20))) & (op2(op2(e24, e20), e25) = op2(e24, op2(e20, e25))) & (op2(op2(e24, e20), e24) = op2(e24, op2(e20, e24))) & (op2(op2(e24, e20), e23) = op2(e24, op2(e20, e23))) & (op2(op2(e24, e20), e22) = op2(e24, op2(e20, e22))) & (op2(op2(e24, e20), e21) = op2(e24, op2(e20, e21))) & (op2(op2(e24, e20), e20) = op2(e24, op2(e20, e20))) & (op2(op2(e23, e25), e25) = op2(e23, op2(e25, e25))) & (op2(op2(e23, e25), e24) = op2(e23, op2(e25, e24))) & (op2(op2(e23, e25), e23) = op2(e23, op2(e25, e23))) & (op2(op2(e23, e25), e22) = op2(e23, op2(e25, e22))) & (op2(op2(e23, e25), e21) = op2(e23, op2(e25, e21))) & (op2(op2(e23, e25), e20) = op2(e23, op2(e25, e20))) & (op2(op2(e23, e24), e25) = op2(e23, op2(e24, e25))) & (op2(op2(e23, e24), e24) = op2(e23, op2(e24, e24))) & (op2(op2(e23, e24), e23) = op2(e23, op2(e24, e23))) & (op2(op2(e23, e24), e22) = op2(e23, op2(e24, e22))) & (op2(op2(e23, e24), e21) = op2(e23, op2(e24, e21))) & (op2(op2(e23, e24), e20) = op2(e23, op2(e24, e20))) & (op2(op2(e23, e23), e25) = op2(e23, op2(e23, e25))) & (op2(op2(e23, e23), e24) = op2(e23, op2(e23, e24))) & (op2(op2(e23, e23), e23) = op2(e23, op2(e23, e23))) & (op2(op2(e23, e23), e22) = op2(e23, op2(e23, e22))) & (op2(op2(e23, e23), e21) = op2(e23, op2(e23, e21))) & (op2(op2(e23, e23), e20) = op2(e23, op2(e23, e20))) & (op2(op2(e23, e22), e25) = op2(e23, op2(e22, e25))) & (op2(op2(e23, e22), e24) = op2(e23, op2(e22, e24))) & (op2(op2(e23, e22), e23) = op2(e23, op2(e22, e23))) & (op2(op2(e23, e22), e22) = op2(e23, op2(e22, e22))) & (op2(op2(e23, e22), e21) = op2(e23, op2(e22, e21))) & (op2(op2(e23, e22), e20) = op2(e23, op2(e22, e20))) & (op2(op2(e23, e21), e25) = op2(e23, op2(e21, e25))) & (op2(op2(e23, e21), e24) = op2(e23, op2(e21, e24))) & (op2(op2(e23, e21), e23) = op2(e23, op2(e21, e23))) & (op2(op2(e23, e21), e22) = op2(e23, op2(e21, e22))) & (op2(op2(e23, e21), e21) = op2(e23, op2(e21, e21))) & (op2(op2(e23, e21), e20) = op2(e23, op2(e21, e20))) & (op2(op2(e23, e20), e25) = op2(e23, op2(e20, e25))) & (op2(op2(e23, e20), e24) = op2(e23, op2(e20, e24))) & (op2(op2(e23, e20), e23) = op2(e23, op2(e20, e23))) & (op2(op2(e23, e20), e22) = op2(e23, op2(e20, e22))) & (op2(op2(e23, e20), e21) = op2(e23, op2(e20, e21))) & (op2(op2(e23, e20), e20) = op2(e23, op2(e20, e20))) & (op2(op2(e22, e25), e25) = op2(e22, op2(e25, e25))) & (op2(op2(e22, e25), e24) = op2(e22, op2(e25, e24))) & (op2(op2(e22, e25), e23) = op2(e22, op2(e25, e23))) & (op2(op2(e22, e25), e22) = op2(e22, op2(e25, e22))) & (op2(op2(e22, e25), e21) = op2(e22, op2(e25, e21))) & (op2(op2(e22, e25), e20) = op2(e22, op2(e25, e20))) & (op2(op2(e22, e24), e25) = op2(e22, op2(e24, e25))) & (op2(op2(e22, e24), e24) = op2(e22, op2(e24, e24))) & (op2(op2(e22, e24), e23) = op2(e22, op2(e24, e23))) & (op2(op2(e22, e24), e22) = op2(e22, op2(e24, e22))) & (op2(op2(e22, e24), e21) = op2(e22, op2(e24, e21))) & (op2(op2(e22, e24), e20) = op2(e22, op2(e24, e20))) & (op2(op2(e22, e23), e25) = op2(e22, op2(e23, e25))) & (op2(op2(e22, e23), e24) = op2(e22, op2(e23, e24))) & (op2(op2(e22, e23), e23) = op2(e22, op2(e23, e23))) & (op2(op2(e22, e23), e22) = op2(e22, op2(e23, e22))) & (op2(op2(e22, e23), e21) = op2(e22, op2(e23, e21))) & (op2(op2(e22, e23), e20) = op2(e22, op2(e23, e20))) & (op2(op2(e22, e22), e25) = op2(e22, op2(e22, e25))) & (op2(op2(e22, e22), e24) = op2(e22, op2(e22, e24))) & (op2(op2(e22, e22), e23) = op2(e22, op2(e22, e23))) & (op2(op2(e22, e22), e22) = op2(e22, op2(e22, e22))) & (op2(op2(e22, e22), e21) = op2(e22, op2(e22, e21))) & (op2(op2(e22, e22), e20) = op2(e22, op2(e22, e20))) & (op2(op2(e22, e21), e25) = op2(e22, op2(e21, e25))) & (op2(op2(e22, e21), e24) = op2(e22, op2(e21, e24))) & (op2(op2(e22, e21), e23) = op2(e22, op2(e21, e23))) & (op2(op2(e22, e21), e22) = op2(e22, op2(e21, e22))) & (op2(op2(e22, e21), e21) = op2(e22, op2(e21, e21))) & (op2(op2(e22, e21), e20) = op2(e22, op2(e21, e20))) & (op2(op2(e22, e20), e25) = op2(e22, op2(e20, e25))) & (op2(op2(e22, e20), e24) = op2(e22, op2(e20, e24))) & (op2(op2(e22, e20), e23) = op2(e22, op2(e20, e23))) & (op2(op2(e22, e20), e22) = op2(e22, op2(e20, e22))) & (op2(op2(e22, e20), e21) = op2(e22, op2(e20, e21))) & (op2(op2(e22, e20), e20) = op2(e22, op2(e20, e20))) & (op2(op2(e21, e25), e25) = op2(e21, op2(e25, e25))) & (op2(op2(e21, e25), e24) = op2(e21, op2(e25, e24))) & (op2(op2(e21, e25), e23) = op2(e21, op2(e25, e23))) & (op2(op2(e21, e25), e22) = op2(e21, op2(e25, e22))) & (op2(op2(e21, e25), e21) = op2(e21, op2(e25, e21))) & (op2(op2(e21, e25), e20) = op2(e21, op2(e25, e20))) & (op2(op2(e21, e24), e25) = op2(e21, op2(e24, e25))) & (op2(op2(e21, e24), e24) = op2(e21, op2(e24, e24))) & (op2(op2(e21, e24), e23) = op2(e21, op2(e24, e23))) & (op2(op2(e21, e24), e22) = op2(e21, op2(e24, e22))) & (op2(op2(e21, e24), e21) = op2(e21, op2(e24, e21))) & (op2(op2(e21, e24), e20) = op2(e21, op2(e24, e20))) & (op2(op2(e21, e23), e25) = op2(e21, op2(e23, e25))) & (op2(op2(e21, e23), e24) = op2(e21, op2(e23, e24))) & (op2(op2(e21, e23), e23) = op2(e21, op2(e23, e23))) & (op2(op2(e21, e23), e22) = op2(e21, op2(e23, e22))) & (op2(op2(e21, e23), e21) = op2(e21, op2(e23, e21))) & (op2(op2(e21, e23), e20) = op2(e21, op2(e23, e20))) & (op2(op2(e21, e22), e25) = op2(e21, op2(e22, e25))) & (op2(op2(e21, e22), e24) = op2(e21, op2(e22, e24))) & (op2(op2(e21, e22), e23) = op2(e21, op2(e22, e23))) & (op2(op2(e21, e22), e22) = op2(e21, op2(e22, e22))) & (op2(op2(e21, e22), e21) = op2(e21, op2(e22, e21))) & (op2(op2(e21, e22), e20) = op2(e21, op2(e22, e20))) & (op2(op2(e21, e21), e25) = op2(e21, op2(e21, e25))) & (op2(op2(e21, e21), e24) = op2(e21, op2(e21, e24))) & (op2(op2(e21, e21), e23) = op2(e21, op2(e21, e23))) & (op2(op2(e21, e21), e22) = op2(e21, op2(e21, e22))) & (op2(op2(e21, e21), e21) = op2(e21, op2(e21, e21))) & (op2(op2(e21, e21), e20) = op2(e21, op2(e21, e20))) & (op2(op2(e21, e20), e25) = op2(e21, op2(e20, e25))) & (op2(op2(e21, e20), e24) = op2(e21, op2(e20, e24))) & (op2(op2(e21, e20), e23) = op2(e21, op2(e20, e23))) & (op2(op2(e21, e20), e22) = op2(e21, op2(e20, e22))) & (op2(op2(e21, e20), e21) = op2(e21, op2(e20, e21))) & (op2(op2(e21, e20), e20) = op2(e21, op2(e20, e20))) & (op2(op2(e20, e25), e25) = op2(e20, op2(e25, e25))) & (op2(op2(e20, e25), e24) = op2(e20, op2(e25, e24))) & (op2(op2(e20, e25), e23) = op2(e20, op2(e25, e23))) & (op2(op2(e20, e25), e22) = op2(e20, op2(e25, e22))) & (op2(op2(e20, e25), e21) = op2(e20, op2(e25, e21))) & (op2(op2(e20, e25), e20) = op2(e20, op2(e25, e20))) & (op2(op2(e20, e24), e25) = op2(e20, op2(e24, e25))) & (op2(op2(e20, e24), e24) = op2(e20, op2(e24, e24))) & (op2(op2(e20, e24), e23) = op2(e20, op2(e24, e23))) & (op2(op2(e20, e24), e22) = op2(e20, op2(e24, e22))) & (op2(op2(e20, e24), e21) = op2(e20, op2(e24, e21))) & (op2(op2(e20, e24), e20) = op2(e20, op2(e24, e20))) & (op2(op2(e20, e23), e25) = op2(e20, op2(e23, e25))) & (op2(op2(e20, e23), e24) = op2(e20, op2(e23, e24))) & (op2(op2(e20, e23), e23) = op2(e20, op2(e23, e23))) & (op2(op2(e20, e23), e22) = op2(e20, op2(e23, e22))) & (op2(op2(e20, e23), e21) = op2(e20, op2(e23, e21))) & (op2(op2(e20, e23), e20) = op2(e20, op2(e23, e20))) & (op2(op2(e20, e22), e25) = op2(e20, op2(e22, e25))) & (op2(op2(e20, e22), e24) = op2(e20, op2(e22, e24))) & (op2(op2(e20, e22), e23) = op2(e20, op2(e22, e23))) & (op2(op2(e20, e22), e22) = op2(e20, op2(e22, e22))) & (op2(op2(e20, e22), e21) = op2(e20, op2(e22, e21))) & (op2(op2(e20, e22), e20) = op2(e20, op2(e22, e20))) & (op2(op2(e20, e21), e25) = op2(e20, op2(e21, e25))) & (op2(op2(e20, e21), e24) = op2(e20, op2(e21, e24))) & (op2(op2(e20, e21), e23) = op2(e20, op2(e21, e23))) & (op2(op2(e20, e21), e22) = op2(e20, op2(e21, e22))) & (op2(op2(e20, e21), e21) = op2(e20, op2(e21, e21))) & (op2(op2(e20, e21), e20) = op2(e20, op2(e21, e20))) & (op2(op2(e20, e20), e25) = op2(e20, op2(e20, e25))) & (op2(op2(e20, e20), e24) = op2(e20, op2(e20, e24))) & (op2(op2(e20, e20), e23) = op2(e20, op2(e20, e23))) & (op2(op2(e20, e20), e22) = op2(e20, op2(e20, e22))) & (op2(op2(e20, e20), e21) = op2(e20, op2(e20, e21))) & (op2(op2(e20, e20), e20) = op2(e20, op2(e20, e20)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG026+1.p', ax6)).
fof(f19118, plain, (~ spl30_261 | ~ spl30_3 | ~ spl30_558 | spl30_731 | ~ spl30_732), inference(avatar_split_clause, [], [f19102, f6305, f6301, f5189, f1549, f2624])).
fof(f1549, plain, (spl30_3 <=> (e12 = op1(e15, e15))), introduced(avatar_definition, [new_symbols(naming, [spl30_3])])).
fof(f6301, plain, (spl30_731 <=> (h5(op1(e15, e15)) = op2(h5(e15), h5(e15)))), introduced(avatar_definition, [new_symbols(naming, [spl30_731])])).
fof(f19102, plain, (~ (e22 = op2(e25, e25)) | (~ spl30_3 | ~ spl30_558 | spl30_731 | ~ spl30_732)), inference(forward_demodulation, [], [f19101, f5190])).
fof(f19101, plain, (~ (op2(e25, e25) = h5(e12)) | (~ spl30_3 | spl30_731 | ~ spl30_732)), inference(forward_demodulation, [], [f19100, f1551])).
fof(f1551, plain, ((e12 = op1(e15, e15)) | ~ spl30_3), inference(avatar_component_clause, [], [f1549])).
fof(f19100, plain, (~ (op2(e25, e25) = h5(op1(e15, e15))) | (spl30_731 | ~ spl30_732)), inference(forward_demodulation, [], [f6303, f6306])).
fof(f6303, plain, (~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | spl30_731), inference(avatar_component_clause, [], [f6301])).
fof(f19049, plain, (~ spl30_77 | ~ spl30_335 | ~ spl30_553 | spl30_730 | ~ spl30_732), inference(avatar_contradiction_clause, [], [f19048])).
fof(f19048, plain, ($false | (~ spl30_77 | ~ spl30_335 | ~ spl30_553 | spl30_730 | ~ spl30_732)), inference(subsumption_resolution, [], [f19047, f2934])).
fof(f2934, plain, ((e24 = op2(e23, e25)) | ~ spl30_335), inference(avatar_component_clause, [], [f2932])).
fof(f2932, plain, (spl30_335 <=> (e24 = op2(e23, e25))), introduced(avatar_definition, [new_symbols(naming, [spl30_335])])).
fof(f19047, plain, (~ (e24 = op2(e23, e25)) | (~ spl30_77 | ~ spl30_553 | spl30_730 | ~ spl30_732)), inference(forward_demodulation, [], [f19046, f1312])).
fof(f19046, plain, (~ (op2(e23, e25) = h5(e14)) | (~ spl30_77 | ~ spl30_553 | spl30_730 | ~ spl30_732)), inference(forward_demodulation, [], [f19045, f1859])).
fof(f19045, plain, (~ (op2(e23, e25) = h5(op1(e13, e15))) | (~ spl30_553 | spl30_730 | ~ spl30_732)), inference(forward_demodulation, [], [f19044, f5163])).
fof(f19044, plain, (~ (h5(op1(e13, e15)) = op2(h5(e13), e25)) | (spl30_730 | ~ spl30_732)), inference(forward_demodulation, [], [f6299, f6306])).
fof(f6299, plain, (~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | spl30_730), inference(avatar_component_clause, [], [f6297])).
fof(f6297, plain, (spl30_730 <=> (h5(op1(e13, e15)) = op2(h5(e13), h5(e15)))), introduced(avatar_definition, [new_symbols(naming, [spl30_730])])).
fof(f18983, plain, (~ spl30_87 | ~ spl30_345 | ~ spl30_553 | ~ spl30_558 | spl30_729), inference(avatar_contradiction_clause, [], [f18982])).
fof(f18982, plain, ($false | (~ spl30_87 | ~ spl30_345 | ~ spl30_553 | ~ spl30_558 | spl30_729)), inference(subsumption_resolution, [], [f18981, f2976])).
fof(f2976, plain, ((e22 = op2(e23, e23)) | ~ spl30_345), inference(avatar_component_clause, [], [f2974])).
fof(f2974, plain, (spl30_345 <=> (e22 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl30_345])])).
fof(f18981, plain, (~ (e22 = op2(e23, e23)) | (~ spl30_87 | ~ spl30_553 | ~ spl30_558 | spl30_729)), inference(forward_demodulation, [], [f18980, f5190])).
fof(f18980, plain, (~ (op2(e23, e23) = h5(e12)) | (~ spl30_87 | ~ spl30_553 | spl30_729)), inference(forward_demodulation, [], [f18979, f1901])).
fof(f1901, plain, ((e12 = op1(e13, e13)) | ~ spl30_87), inference(avatar_component_clause, [], [f1899])).
fof(f1899, plain, (spl30_87 <=> (e12 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl30_87])])).
fof(f18979, plain, (~ (op2(e23, e23) = h5(op1(e13, e13))) | (~ spl30_553 | spl30_729)), inference(forward_demodulation, [], [f6295, f5163])).
fof(f6295, plain, (~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | spl30_729), inference(avatar_component_clause, [], [f6293])).
fof(f6293, plain, (spl30_729 <=> (h5(op1(e13, e13)) = op2(h5(e13), h5(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl30_729])])).
fof(f18837, plain, (~ spl30_110 | ~ spl30_368 | ~ spl30_558 | spl30_728 | ~ spl30_732), inference(avatar_contradiction_clause, [], [f18836])).
fof(f18836, plain, ($false | (~ spl30_110 | ~ spl30_368 | ~ spl30_558 | spl30_728 | ~ spl30_732)), inference(subsumption_resolution, [], [f18835, f3072])).
fof(f3072, plain, ((e21 = op2(e22, e25)) | ~ spl30_368), inference(avatar_component_clause, [], [f3070])).
fof(f3070, plain, (spl30_368 <=> (e21 = op2(e22, e25))), introduced(avatar_definition, [new_symbols(naming, [spl30_368])])).
fof(f18835, plain, (~ (e21 = op2(e22, e25)) | (~ spl30_110 | ~ spl30_558 | spl30_728 | ~ spl30_732)), inference(forward_demodulation, [], [f18834, f4941])).
fof(f18834, plain, (~ (op2(e22, e25) = h5(e11)) | (~ spl30_110 | ~ spl30_558 | spl30_728 | ~ spl30_732)), inference(forward_demodulation, [], [f18833, f1997])).
fof(f18833, plain, (~ (op2(e22, e25) = h5(op1(e12, e15))) | (~ spl30_558 | spl30_728 | ~ spl30_732)), inference(forward_demodulation, [], [f18832, f5190])).
fof(f18832, plain, (~ (h5(op1(e12, e15)) = op2(h5(e12), e25)) | (spl30_728 | ~ spl30_732)), inference(forward_demodulation, [], [f6291, f6306])).
fof(f6291, plain, (~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | spl30_728), inference(avatar_component_clause, [], [f6289])).
fof(f6289, plain, (spl30_728 <=> (h5(op1(e12, e15)) = op2(h5(e12), h5(e15)))), introduced(avatar_definition, [new_symbols(naming, [spl30_728])])).
fof(f18614, plain, (~ spl30_191 | spl30_221 | ~ spl30_229), inference(avatar_contradiction_clause, [], [f18613])).
fof(f18613, plain, ($false | (~ spl30_191 | spl30_221 | ~ spl30_229)), inference(subsumption_resolution, [], [f18612, f2458])).
fof(f2458, plain, (~ (e14 = unit1) | spl30_221), inference(avatar_component_clause, [], [f2457])).
fof(f2457, plain, (spl30_221 <=> (e14 = unit1)), introduced(avatar_definition, [new_symbols(naming, [spl30_221])])).
fof(f18612, plain, ((e14 = unit1) | (~ spl30_191 | ~ spl30_229)), inference(forward_demodulation, [], [f18601, f2334])).
fof(f18601, plain, ((op1(e10, e14) = unit1) | ~ spl30_229), inference(backward_demodulation, [], [f372, f2493])).
fof(f2493, plain, ((e10 = inv1(e14)) | ~ spl30_229), inference(avatar_component_clause, [], [f2491])).
fof(f2491, plain, (spl30_229 <=> (e10 = inv1(e14))), introduced(avatar_definition, [new_symbols(naming, [spl30_229])])).
fof(f372, plain, (unit1 = op1(inv1(e14), e14)), inference(cnf_transformation, [], [f4])).
fof(f4, plain, (((e15 = inv1(e15)) | (e14 = inv1(e15)) | (e13 = inv1(e15)) | (e12 = inv1(e15)) | (e11 = inv1(e15)) | (e10 = inv1(e15))) & ((e15 = inv1(e14)) | (e14 = inv1(e14)) | (e13 = inv1(e14)) | (e12 = inv1(e14)) | (e11 = inv1(e14)) | (e10 = inv1(e14))) & ((e15 = inv1(e13)) | (e14 = inv1(e13)) | (e13 = inv1(e13)) | (e12 = inv1(e13)) | (e11 = inv1(e13)) | (e10 = inv1(e13))) & ((e15 = inv1(e12)) | (e14 = inv1(e12)) | (e13 = inv1(e12)) | (e12 = inv1(e12)) | (e11 = inv1(e12)) | (e10 = inv1(e12))) & ((e15 = inv1(e11)) | (e14 = inv1(e11)) | (e13 = inv1(e11)) | (e12 = inv1(e11)) | (e11 = inv1(e11)) | (e10 = inv1(e11))) & ((e15 = inv1(e10)) | (e14 = inv1(e10)) | (e13 = inv1(e10)) | (e12 = inv1(e10)) | (e11 = inv1(e10)) | (e10 = inv1(e10))) & (unit1 = op1(inv1(e15), e15)) & (unit1 = op1(e15, inv1(e15))) & (unit1 = op1(inv1(e14), e14)) & (unit1 = op1(e14, inv1(e14))) & (unit1 = op1(inv1(e13), e13)) & (unit1 = op1(e13, inv1(e13))) & (unit1 = op1(inv1(e12), e12)) & (unit1 = op1(e12, inv1(e12))) & (unit1 = op1(inv1(e11), e11)) & (unit1 = op1(e11, inv1(e11))) & (unit1 = op1(inv1(e10), e10)) & (unit1 = op1(e10, inv1(e10)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG026+1.p', ax4)).
fof(f18302, plain, (spl30_167 | ~ spl30_153 | ~ spl30_169 | ~ spl30_191), inference(avatar_split_clause, [], [f18301, f2332, f2241, f2174, f2232])).
fof(f2174, plain, (spl30_153 <=> (e12 = op1(e11, e14))), introduced(avatar_definition, [new_symbols(naming, [spl30_153])])).
fof(f18301, plain, ((e14 = op1(e11, e12)) | (~ spl30_153 | ~ spl30_169 | ~ spl30_191)), inference(forward_demodulation, [], [f18300, f2334])).
fof(f18300, plain, ((op1(e10, e14) = op1(e11, e12)) | (~ spl30_153 | ~ spl30_169)), inference(forward_demodulation, [], [f17544, f2243])).
fof(f17544, plain, ((op1(e11, e12) = op1(op1(e11, e11), e14)) | ~ spl30_153), inference(forward_demodulation, [], [f180, f2176])).
fof(f2176, plain, ((e12 = op1(e11, e14)) | ~ spl30_153), inference(avatar_component_clause, [], [f2174])).
fof(f180, plain, (op1(op1(e11, e11), e14) = op1(e11, op1(e11, e14))), inference(cnf_transformation, [], [f2])).
fof(f2, plain, ((op1(op1(e15, e15), e15) = op1(e15, op1(e15, e15))) & (op1(op1(e15, e15), e14) = op1(e15, op1(e15, e14))) & (op1(op1(e15, e15), e13) = op1(e15, op1(e15, e13))) & (op1(op1(e15, e15), e12) = op1(e15, op1(e15, e12))) & (op1(op1(e15, e15), e11) = op1(e15, op1(e15, e11))) & (op1(op1(e15, e15), e10) = op1(e15, op1(e15, e10))) & (op1(op1(e15, e14), e15) = op1(e15, op1(e14, e15))) & (op1(op1(e15, e14), e14) = op1(e15, op1(e14, e14))) & (op1(op1(e15, e14), e13) = op1(e15, op1(e14, e13))) & (op1(op1(e15, e14), e12) = op1(e15, op1(e14, e12))) & (op1(op1(e15, e14), e11) = op1(e15, op1(e14, e11))) & (op1(op1(e15, e14), e10) = op1(e15, op1(e14, e10))) & (op1(op1(e15, e13), e15) = op1(e15, op1(e13, e15))) & (op1(op1(e15, e13), e14) = op1(e15, op1(e13, e14))) & (op1(op1(e15, e13), e13) = op1(e15, op1(e13, e13))) & (op1(op1(e15, e13), e12) = op1(e15, op1(e13, e12))) & (op1(op1(e15, e13), e11) = op1(e15, op1(e13, e11))) & (op1(op1(e15, e13), e10) = op1(e15, op1(e13, e10))) & (op1(op1(e15, e12), e15) = op1(e15, op1(e12, e15))) & (op1(op1(e15, e12), e14) = op1(e15, op1(e12, e14))) & (op1(op1(e15, e12), e13) = op1(e15, op1(e12, e13))) & (op1(op1(e15, e12), e12) = op1(e15, op1(e12, e12))) & (op1(op1(e15, e12), e11) = op1(e15, op1(e12, e11))) & (op1(op1(e15, e12), e10) = op1(e15, op1(e12, e10))) & (op1(op1(e15, e11), e15) = op1(e15, op1(e11, e15))) & (op1(op1(e15, e11), e14) = op1(e15, op1(e11, e14))) & (op1(op1(e15, e11), e13) = op1(e15, op1(e11, e13))) & (op1(op1(e15, e11), e12) = op1(e15, op1(e11, e12))) & (op1(op1(e15, e11), e11) = op1(e15, op1(e11, e11))) & (op1(op1(e15, e11), e10) = op1(e15, op1(e11, e10))) & (op1(op1(e15, e10), e15) = op1(e15, op1(e10, e15))) & (op1(op1(e15, e10), e14) = op1(e15, op1(e10, e14))) & (op1(op1(e15, e10), e13) = op1(e15, op1(e10, e13))) & (op1(op1(e15, e10), e12) = op1(e15, op1(e10, e12))) & (op1(op1(e15, e10), e11) = op1(e15, op1(e10, e11))) & (op1(op1(e15, e10), e10) = op1(e15, op1(e10, e10))) & (op1(op1(e14, e15), e15) = op1(e14, op1(e15, e15))) & (op1(op1(e14, e15), e14) = op1(e14, op1(e15, e14))) & (op1(op1(e14, e15), e13) = op1(e14, op1(e15, e13))) & (op1(op1(e14, e15), e12) = op1(e14, op1(e15, e12))) & (op1(op1(e14, e15), e11) = op1(e14, op1(e15, e11))) & (op1(op1(e14, e15), e10) = op1(e14, op1(e15, e10))) & (op1(op1(e14, e14), e15) = op1(e14, op1(e14, e15))) & (op1(op1(e14, e14), e14) = op1(e14, op1(e14, e14))) & (op1(op1(e14, e14), e13) = op1(e14, op1(e14, e13))) & (op1(op1(e14, e14), e12) = op1(e14, op1(e14, e12))) & (op1(op1(e14, e14), e11) = op1(e14, op1(e14, e11))) & (op1(op1(e14, e14), e10) = op1(e14, op1(e14, e10))) & (op1(op1(e14, e13), e15) = op1(e14, op1(e13, e15))) & (op1(op1(e14, e13), e14) = op1(e14, op1(e13, e14))) & (op1(op1(e14, e13), e13) = op1(e14, op1(e13, e13))) & (op1(op1(e14, e13), e12) = op1(e14, op1(e13, e12))) & (op1(op1(e14, e13), e11) = op1(e14, op1(e13, e11))) & (op1(op1(e14, e13), e10) = op1(e14, op1(e13, e10))) & (op1(op1(e14, e12), e15) = op1(e14, op1(e12, e15))) & (op1(op1(e14, e12), e14) = op1(e14, op1(e12, e14))) & (op1(op1(e14, e12), e13) = op1(e14, op1(e12, e13))) & (op1(op1(e14, e12), e12) = op1(e14, op1(e12, e12))) & (op1(op1(e14, e12), e11) = op1(e14, op1(e12, e11))) & (op1(op1(e14, e12), e10) = op1(e14, op1(e12, e10))) & (op1(op1(e14, e11), e15) = op1(e14, op1(e11, e15))) & (op1(op1(e14, e11), e14) = op1(e14, op1(e11, e14))) & (op1(op1(e14, e11), e13) = op1(e14, op1(e11, e13))) & (op1(op1(e14, e11), e12) = op1(e14, op1(e11, e12))) & (op1(op1(e14, e11), e11) = op1(e14, op1(e11, e11))) & (op1(op1(e14, e11), e10) = op1(e14, op1(e11, e10))) & (op1(op1(e14, e10), e15) = op1(e14, op1(e10, e15))) & (op1(op1(e14, e10), e14) = op1(e14, op1(e10, e14))) & (op1(op1(e14, e10), e13) = op1(e14, op1(e10, e13))) & (op1(op1(e14, e10), e12) = op1(e14, op1(e10, e12))) & (op1(op1(e14, e10), e11) = op1(e14, op1(e10, e11))) & (op1(op1(e14, e10), e10) = op1(e14, op1(e10, e10))) & (op1(op1(e13, e15), e15) = op1(e13, op1(e15, e15))) & (op1(op1(e13, e15), e14) = op1(e13, op1(e15, e14))) & (op1(op1(e13, e15), e13) = op1(e13, op1(e15, e13))) & (op1(op1(e13, e15), e12) = op1(e13, op1(e15, e12))) & (op1(op1(e13, e15), e11) = op1(e13, op1(e15, e11))) & (op1(op1(e13, e15), e10) = op1(e13, op1(e15, e10))) & (op1(op1(e13, e14), e15) = op1(e13, op1(e14, e15))) & (op1(op1(e13, e14), e14) = op1(e13, op1(e14, e14))) & (op1(op1(e13, e14), e13) = op1(e13, op1(e14, e13))) & (op1(op1(e13, e14), e12) = op1(e13, op1(e14, e12))) & (op1(op1(e13, e14), e11) = op1(e13, op1(e14, e11))) & (op1(op1(e13, e14), e10) = op1(e13, op1(e14, e10))) & (op1(op1(e13, e13), e15) = op1(e13, op1(e13, e15))) & (op1(op1(e13, e13), e14) = op1(e13, op1(e13, e14))) & (op1(op1(e13, e13), e13) = op1(e13, op1(e13, e13))) & (op1(op1(e13, e13), e12) = op1(e13, op1(e13, e12))) & (op1(op1(e13, e13), e11) = op1(e13, op1(e13, e11))) & (op1(op1(e13, e13), e10) = op1(e13, op1(e13, e10))) & (op1(op1(e13, e12), e15) = op1(e13, op1(e12, e15))) & (op1(op1(e13, e12), e14) = op1(e13, op1(e12, e14))) & (op1(op1(e13, e12), e13) = op1(e13, op1(e12, e13))) & (op1(op1(e13, e12), e12) = op1(e13, op1(e12, e12))) & (op1(op1(e13, e12), e11) = op1(e13, op1(e12, e11))) & (op1(op1(e13, e12), e10) = op1(e13, op1(e12, e10))) & (op1(op1(e13, e11), e15) = op1(e13, op1(e11, e15))) & (op1(op1(e13, e11), e14) = op1(e13, op1(e11, e14))) & (op1(op1(e13, e11), e13) = op1(e13, op1(e11, e13))) & (op1(op1(e13, e11), e12) = op1(e13, op1(e11, e12))) & (op1(op1(e13, e11), e11) = op1(e13, op1(e11, e11))) & (op1(op1(e13, e11), e10) = op1(e13, op1(e11, e10))) & (op1(op1(e13, e10), e15) = op1(e13, op1(e10, e15))) & (op1(op1(e13, e10), e14) = op1(e13, op1(e10, e14))) & (op1(op1(e13, e10), e13) = op1(e13, op1(e10, e13))) & (op1(op1(e13, e10), e12) = op1(e13, op1(e10, e12))) & (op1(op1(e13, e10), e11) = op1(e13, op1(e10, e11))) & (op1(op1(e13, e10), e10) = op1(e13, op1(e10, e10))) & (op1(op1(e12, e15), e15) = op1(e12, op1(e15, e15))) & (op1(op1(e12, e15), e14) = op1(e12, op1(e15, e14))) & (op1(op1(e12, e15), e13) = op1(e12, op1(e15, e13))) & (op1(op1(e12, e15), e12) = op1(e12, op1(e15, e12))) & (op1(op1(e12, e15), e11) = op1(e12, op1(e15, e11))) & (op1(op1(e12, e15), e10) = op1(e12, op1(e15, e10))) & (op1(op1(e12, e14), e15) = op1(e12, op1(e14, e15))) & (op1(op1(e12, e14), e14) = op1(e12, op1(e14, e14))) & (op1(op1(e12, e14), e13) = op1(e12, op1(e14, e13))) & (op1(op1(e12, e14), e12) = op1(e12, op1(e14, e12))) & (op1(op1(e12, e14), e11) = op1(e12, op1(e14, e11))) & (op1(op1(e12, e14), e10) = op1(e12, op1(e14, e10))) & (op1(op1(e12, e13), e15) = op1(e12, op1(e13, e15))) & (op1(op1(e12, e13), e14) = op1(e12, op1(e13, e14))) & (op1(op1(e12, e13), e13) = op1(e12, op1(e13, e13))) & (op1(op1(e12, e13), e12) = op1(e12, op1(e13, e12))) & (op1(op1(e12, e13), e11) = op1(e12, op1(e13, e11))) & (op1(op1(e12, e13), e10) = op1(e12, op1(e13, e10))) & (op1(op1(e12, e12), e15) = op1(e12, op1(e12, e15))) & (op1(op1(e12, e12), e14) = op1(e12, op1(e12, e14))) & (op1(op1(e12, e12), e13) = op1(e12, op1(e12, e13))) & (op1(op1(e12, e12), e12) = op1(e12, op1(e12, e12))) & (op1(op1(e12, e12), e11) = op1(e12, op1(e12, e11))) & (op1(op1(e12, e12), e10) = op1(e12, op1(e12, e10))) & (op1(op1(e12, e11), e15) = op1(e12, op1(e11, e15))) & (op1(op1(e12, e11), e14) = op1(e12, op1(e11, e14))) & (op1(op1(e12, e11), e13) = op1(e12, op1(e11, e13))) & (op1(op1(e12, e11), e12) = op1(e12, op1(e11, e12))) & (op1(op1(e12, e11), e11) = op1(e12, op1(e11, e11))) & (op1(op1(e12, e11), e10) = op1(e12, op1(e11, e10))) & (op1(op1(e12, e10), e15) = op1(e12, op1(e10, e15))) & (op1(op1(e12, e10), e14) = op1(e12, op1(e10, e14))) & (op1(op1(e12, e10), e13) = op1(e12, op1(e10, e13))) & (op1(op1(e12, e10), e12) = op1(e12, op1(e10, e12))) & (op1(op1(e12, e10), e11) = op1(e12, op1(e10, e11))) & (op1(op1(e12, e10), e10) = op1(e12, op1(e10, e10))) & (op1(op1(e11, e15), e15) = op1(e11, op1(e15, e15))) & (op1(op1(e11, e15), e14) = op1(e11, op1(e15, e14))) & (op1(op1(e11, e15), e13) = op1(e11, op1(e15, e13))) & (op1(op1(e11, e15), e12) = op1(e11, op1(e15, e12))) & (op1(op1(e11, e15), e11) = op1(e11, op1(e15, e11))) & (op1(op1(e11, e15), e10) = op1(e11, op1(e15, e10))) & (op1(op1(e11, e14), e15) = op1(e11, op1(e14, e15))) & (op1(op1(e11, e14), e14) = op1(e11, op1(e14, e14))) & (op1(op1(e11, e14), e13) = op1(e11, op1(e14, e13))) & (op1(op1(e11, e14), e12) = op1(e11, op1(e14, e12))) & (op1(op1(e11, e14), e11) = op1(e11, op1(e14, e11))) & (op1(op1(e11, e14), e10) = op1(e11, op1(e14, e10))) & (op1(op1(e11, e13), e15) = op1(e11, op1(e13, e15))) & (op1(op1(e11, e13), e14) = op1(e11, op1(e13, e14))) & (op1(op1(e11, e13), e13) = op1(e11, op1(e13, e13))) & (op1(op1(e11, e13), e12) = op1(e11, op1(e13, e12))) & (op1(op1(e11, e13), e11) = op1(e11, op1(e13, e11))) & (op1(op1(e11, e13), e10) = op1(e11, op1(e13, e10))) & (op1(op1(e11, e12), e15) = op1(e11, op1(e12, e15))) & (op1(op1(e11, e12), e14) = op1(e11, op1(e12, e14))) & (op1(op1(e11, e12), e13) = op1(e11, op1(e12, e13))) & (op1(op1(e11, e12), e12) = op1(e11, op1(e12, e12))) & (op1(op1(e11, e12), e11) = op1(e11, op1(e12, e11))) & (op1(op1(e11, e12), e10) = op1(e11, op1(e12, e10))) & (op1(op1(e11, e11), e15) = op1(e11, op1(e11, e15))) & (op1(op1(e11, e11), e14) = op1(e11, op1(e11, e14))) & (op1(op1(e11, e11), e13) = op1(e11, op1(e11, e13))) & (op1(op1(e11, e11), e12) = op1(e11, op1(e11, e12))) & (op1(op1(e11, e11), e11) = op1(e11, op1(e11, e11))) & (op1(op1(e11, e11), e10) = op1(e11, op1(e11, e10))) & (op1(op1(e11, e10), e15) = op1(e11, op1(e10, e15))) & (op1(op1(e11, e10), e14) = op1(e11, op1(e10, e14))) & (op1(op1(e11, e10), e13) = op1(e11, op1(e10, e13))) & (op1(op1(e11, e10), e12) = op1(e11, op1(e10, e12))) & (op1(op1(e11, e10), e11) = op1(e11, op1(e10, e11))) & (op1(op1(e11, e10), e10) = op1(e11, op1(e10, e10))) & (op1(op1(e10, e15), e15) = op1(e10, op1(e15, e15))) & (op1(op1(e10, e15), e14) = op1(e10, op1(e15, e14))) & (op1(op1(e10, e15), e13) = op1(e10, op1(e15, e13))) & (op1(op1(e10, e15), e12) = op1(e10, op1(e15, e12))) & (op1(op1(e10, e15), e11) = op1(e10, op1(e15, e11))) & (op1(op1(e10, e15), e10) = op1(e10, op1(e15, e10))) & (op1(op1(e10, e14), e15) = op1(e10, op1(e14, e15))) & (op1(op1(e10, e14), e14) = op1(e10, op1(e14, e14))) & (op1(op1(e10, e14), e13) = op1(e10, op1(e14, e13))) & (op1(op1(e10, e14), e12) = op1(e10, op1(e14, e12))) & (op1(op1(e10, e14), e11) = op1(e10, op1(e14, e11))) & (op1(op1(e10, e14), e10) = op1(e10, op1(e14, e10))) & (op1(op1(e10, e13), e15) = op1(e10, op1(e13, e15))) & (op1(op1(e10, e13), e14) = op1(e10, op1(e13, e14))) & (op1(op1(e10, e13), e13) = op1(e10, op1(e13, e13))) & (op1(op1(e10, e13), e12) = op1(e10, op1(e13, e12))) & (op1(op1(e10, e13), e11) = op1(e10, op1(e13, e11))) & (op1(op1(e10, e13), e10) = op1(e10, op1(e13, e10))) & (op1(op1(e10, e12), e15) = op1(e10, op1(e12, e15))) & (op1(op1(e10, e12), e14) = op1(e10, op1(e12, e14))) & (op1(op1(e10, e12), e13) = op1(e10, op1(e12, e13))) & (op1(op1(e10, e12), e12) = op1(e10, op1(e12, e12))) & (op1(op1(e10, e12), e11) = op1(e10, op1(e12, e11))) & (op1(op1(e10, e12), e10) = op1(e10, op1(e12, e10))) & (op1(op1(e10, e11), e15) = op1(e10, op1(e11, e15))) & (op1(op1(e10, e11), e14) = op1(e10, op1(e11, e14))) & (op1(op1(e10, e11), e13) = op1(e10, op1(e11, e13))) & (op1(op1(e10, e11), e12) = op1(e10, op1(e11, e12))) & (op1(op1(e10, e11), e11) = op1(e10, op1(e11, e11))) & (op1(op1(e10, e11), e10) = op1(e10, op1(e11, e10))) & (op1(op1(e10, e10), e15) = op1(e10, op1(e10, e15))) & (op1(op1(e10, e10), e14) = op1(e10, op1(e10, e14))) & (op1(op1(e10, e10), e13) = op1(e10, op1(e10, e13))) & (op1(op1(e10, e10), e12) = op1(e10, op1(e10, e12))) & (op1(op1(e10, e10), e11) = op1(e10, op1(e10, e11))) & (op1(op1(e10, e10), e10) = op1(e10, op1(e10, e10)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG026+1.p', ax2)).
fof(f18257, plain, (spl30_148 | ~ spl30_162 | ~ spl30_169 | ~ spl30_196), inference(avatar_split_clause, [], [f18256, f2353, f2241, f2211, f2153])).
fof(f18256, plain, ((e13 = op1(e11, e15)) | (~ spl30_162 | ~ spl30_169 | ~ spl30_196)), inference(forward_demodulation, [], [f18255, f2355])).
fof(f18255, plain, ((op1(e10, e13) = op1(e11, e15)) | (~ spl30_162 | ~ spl30_169)), inference(forward_demodulation, [], [f17546, f2243])).
fof(f17546, plain, ((op1(e11, e15) = op1(op1(e11, e11), e13)) | ~ spl30_162), inference(forward_demodulation, [], [f179, f2213])).
fof(f179, plain, (op1(op1(e11, e11), e13) = op1(e11, op1(e11, e13))), inference(cnf_transformation, [], [f2])).
fof(f18216, plain, (spl30_199 | ~ spl30_219), inference(avatar_contradiction_clause, [], [f18215])).
fof(f18215, plain, ($false | (spl30_199 | ~ spl30_219)), inference(subsumption_resolution, [], [f18195, f2367])).
fof(f2367, plain, (~ (e10 = op1(e10, e12)) | spl30_199), inference(avatar_component_clause, [], [f2366])).
fof(f2366, plain, (spl30_199 <=> (e10 = op1(e10, e12))), introduced(avatar_definition, [new_symbols(naming, [spl30_199])])).
fof(f18195, plain, ((e10 = op1(e10, e12)) | ~ spl30_219), inference(backward_demodulation, [], [f351, f2451])).
fof(f2451, plain, ((e12 = unit1) | ~ spl30_219), inference(avatar_component_clause, [], [f2449])).
fof(f2449, plain, (spl30_219 <=> (e12 = unit1)), introduced(avatar_definition, [new_symbols(naming, [spl30_219])])).
fof(f351, plain, (e10 = op1(e10, unit1)), inference(cnf_transformation, [], [f3])).
fof(f3, plain, (((e15 = unit1) | (e14 = unit1) | (e13 = unit1) | (e12 = unit1) | (e11 = unit1) | (e10 = unit1)) & (e15 = op1(e15, unit1)) & (e15 = op1(unit1, e15)) & (e14 = op1(e14, unit1)) & (e14 = op1(unit1, e14)) & (e13 = op1(e13, unit1)) & (e13 = op1(unit1, e13)) & (e12 = op1(e12, unit1)) & (e12 = op1(unit1, e12)) & (e11 = op1(e11, unit1)) & (e11 = op1(unit1, e11)) & (e10 = op1(e10, unit1)) & (e10 = op1(unit1, e10))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG026+1.p', ax3)).
fof(f18169, plain, (~ spl30_148 | ~ spl30_406 | ~ spl30_553 | ~ spl30_732 | spl30_744), inference(avatar_contradiction_clause, [], [f18168])).
fof(f18168, plain, ($false | (~ spl30_148 | ~ spl30_406 | ~ spl30_553 | ~ spl30_732 | spl30_744)), inference(subsumption_resolution, [], [f18167, f3230])).
fof(f3230, plain, ((e23 = op2(e21, e25)) | ~ spl30_406), inference(avatar_component_clause, [], [f3228])).
fof(f3228, plain, (spl30_406 <=> (e23 = op2(e21, e25))), introduced(avatar_definition, [new_symbols(naming, [spl30_406])])).
fof(f18167, plain, (~ (e23 = op2(e21, e25)) | (~ spl30_148 | ~ spl30_553 | ~ spl30_732 | spl30_744)), inference(forward_demodulation, [], [f18166, f5163])).
fof(f18166, plain, (~ (op2(e21, e25) = h5(e13)) | (~ spl30_148 | ~ spl30_732 | spl30_744)), inference(forward_demodulation, [], [f18165, f2155])).
fof(f18165, plain, (~ (op2(e21, e25) = h5(op1(e11, e15))) | (~ spl30_732 | spl30_744)), inference(forward_demodulation, [], [f6355, f6306])).
fof(f6355, plain, (~ (h5(op1(e11, e15)) = op2(e21, h5(e15))) | spl30_744), inference(avatar_component_clause, [], [f6353])).
fof(f6353, plain, (spl30_744 <=> (h5(op1(e11, e15)) = op2(e21, h5(e15)))), introduced(avatar_definition, [new_symbols(naming, [spl30_744])])).
fof(f18093, plain, (spl30_219 | ~ spl30_63 | ~ spl30_230), inference(avatar_split_clause, [], [f17788, f2495, f1799, f2449])).
fof(f1799, plain, (spl30_63 <=> (e12 = op1(e14, e11))), introduced(avatar_definition, [new_symbols(naming, [spl30_63])])).
fof(f2495, plain, (spl30_230 <=> (e11 = inv1(e14))), introduced(avatar_definition, [new_symbols(naming, [spl30_230])])).
fof(f17788, plain, ((e12 = unit1) | (~ spl30_63 | ~ spl30_230)), inference(forward_demodulation, [], [f17780, f1801])).
fof(f1801, plain, ((e12 = op1(e14, e11)) | ~ spl30_63), inference(avatar_component_clause, [], [f1799])).
fof(f17780, plain, ((op1(e14, e11) = unit1) | ~ spl30_230), inference(backward_demodulation, [], [f371, f2497])).
fof(f2497, plain, ((e11 = inv1(e14)) | ~ spl30_230), inference(avatar_component_clause, [], [f2495])).
fof(f371, plain, (unit1 = op1(e14, inv1(e14))), inference(cnf_transformation, [], [f4])).
fof(f17998, plain, (spl30_137 | ~ spl30_80 | ~ spl30_121 | ~ spl30_191), inference(avatar_split_clause, [], [f17997, f2332, f2041, f1870, f2107])).
fof(f2107, plain, (spl30_137 <=> (e14 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl30_137])])).
fof(f17997, plain, ((e14 = op1(e12, e11)) | (~ spl30_80 | ~ spl30_121 | ~ spl30_191)), inference(forward_demodulation, [], [f17220, f2334])).
fof(f17220, plain, ((op1(e10, e14) = op1(e12, e11)) | (~ spl30_80 | ~ spl30_121)), inference(forward_demodulation, [], [f11773, f2043])).
fof(f11773, plain, ((op1(e12, e11) = op1(op1(e12, e13), e14)) | ~ spl30_80), inference(forward_demodulation, [], [f228, f1872])).
fof(f228, plain, (op1(op1(e12, e13), e14) = op1(e12, op1(e13, e14))), inference(cnf_transformation, [], [f2])).
fof(f17836, plain, (spl30_117 | ~ spl30_221), inference(avatar_contradiction_clause, [], [f17835])).
fof(f17835, plain, ($false | (spl30_117 | ~ spl30_221)), inference(subsumption_resolution, [], [f17814, f2025])).
fof(f2025, plain, (~ (e12 = op1(e12, e14)) | spl30_117), inference(avatar_component_clause, [], [f2024])).
fof(f2024, plain, (spl30_117 <=> (e12 = op1(e12, e14))), introduced(avatar_definition, [new_symbols(naming, [spl30_117])])).
fof(f17814, plain, ((e12 = op1(e12, e14)) | ~ spl30_221), inference(backward_demodulation, [], [f355, f2459])).
fof(f2459, plain, ((e14 = unit1) | ~ spl30_221), inference(avatar_component_clause, [], [f2457])).
fof(f355, plain, (e12 = op1(e12, unit1)), inference(cnf_transformation, [], [f3])).
fof(f17648, plain, (~ spl30_153 | spl30_745), inference(avatar_contradiction_clause, [], [f17647])).
fof(f17647, plain, ($false | (~ spl30_153 | spl30_745)), inference(trivial_inequality_removal, [], [f17646])).
fof(f17646, plain, (~ (h5(e12) = h5(e12)) | (~ spl30_153 | spl30_745)), inference(forward_demodulation, [], [f6359, f2176])).
fof(f6359, plain, (~ (h5(e12) = h5(op1(e11, e14))) | spl30_745), inference(avatar_component_clause, [], [f6357])).
fof(f6357, plain, (spl30_745 <=> (h5(e12) = h5(op1(e11, e14)))), introduced(avatar_definition, [new_symbols(naming, [spl30_745])])).
fof(f17394, plain, (spl30_146 | ~ spl30_222), inference(avatar_contradiction_clause, [], [f17393])).
fof(f17393, plain, ($false | (spl30_146 | ~ spl30_222)), inference(subsumption_resolution, [], [f17366, f2146])).
fof(f2146, plain, (~ (e11 = op1(e11, e15)) | spl30_146), inference(avatar_component_clause, [], [f2145])).
fof(f2145, plain, (spl30_146 <=> (e11 = op1(e11, e15))), introduced(avatar_definition, [new_symbols(naming, [spl30_146])])).
fof(f17366, plain, ((e11 = op1(e11, e15)) | ~ spl30_222), inference(backward_demodulation, [], [f353, f2463])).
fof(f2463, plain, ((e15 = unit1) | ~ spl30_222), inference(avatar_component_clause, [], [f2461])).
fof(f2461, plain, (spl30_222 <=> (e15 = unit1)), introduced(avatar_definition, [new_symbols(naming, [spl30_222])])).
fof(f353, plain, (e11 = op1(e11, unit1)), inference(cnf_transformation, [], [f3])).
fof(f17347, plain, (~ spl30_50 | spl30_218 | ~ spl30_232), inference(avatar_contradiction_clause, [], [f17346])).
fof(f17346, plain, ($false | (~ spl30_50 | spl30_218 | ~ spl30_232)), inference(subsumption_resolution, [], [f17345, f2446])).
fof(f2446, plain, (~ (e11 = unit1) | spl30_218), inference(avatar_component_clause, [], [f2445])).
fof(f2445, plain, (spl30_218 <=> (e11 = unit1)), introduced(avatar_definition, [new_symbols(naming, [spl30_218])])).
fof(f17345, plain, ((e11 = unit1) | (~ spl30_50 | ~ spl30_232)), inference(forward_demodulation, [], [f17339, f1747])).
fof(f1747, plain, ((e11 = op1(e14, e13)) | ~ spl30_50), inference(avatar_component_clause, [], [f1745])).
fof(f1745, plain, (spl30_50 <=> (e11 = op1(e14, e13))), introduced(avatar_definition, [new_symbols(naming, [spl30_50])])).
fof(f17339, plain, ((op1(e14, e13) = unit1) | ~ spl30_232), inference(backward_demodulation, [], [f371, f2505])).
fof(f2505, plain, ((e13 = inv1(e14)) | ~ spl30_232), inference(avatar_component_clause, [], [f2503])).
fof(f2503, plain, (spl30_232 <=> (e13 = inv1(e14))), introduced(avatar_definition, [new_symbols(naming, [spl30_232])])).
fof(f17321, plain, (spl30_222 | ~ spl30_60 | ~ spl30_245), inference(avatar_split_clause, [], [f17320, f2557, f1786, f2461])).
fof(f1786, plain, (spl30_60 <=> (e15 = op1(e14, e12))), introduced(avatar_definition, [new_symbols(naming, [spl30_60])])).
fof(f2557, plain, (spl30_245 <=> (e14 = inv1(e12))), introduced(avatar_definition, [new_symbols(naming, [spl30_245])])).
fof(f17320, plain, ((e15 = unit1) | (~ spl30_60 | ~ spl30_245)), inference(forward_demodulation, [], [f17312, f1788])).
fof(f1788, plain, ((e15 = op1(e14, e12)) | ~ spl30_60), inference(avatar_component_clause, [], [f1786])).
fof(f17312, plain, ((op1(e14, e12) = unit1) | ~ spl30_245), inference(backward_demodulation, [], [f368, f2559])).
fof(f2559, plain, ((e14 = inv1(e12)) | ~ spl30_245), inference(avatar_component_clause, [], [f2557])).
fof(f368, plain, (unit1 = op1(inv1(e12), e12)), inference(cnf_transformation, [], [f4])).
fof(f17229, plain, (spl30_345 | ~ spl30_581), inference(avatar_split_clause, [], [f17228, f5313, f2974])).
fof(f5313, plain, (spl30_581 <=> (e22 = h4(e13))), introduced(avatar_definition, [new_symbols(naming, [spl30_581])])).
fof(f17228, plain, ((e22 = op2(e23, e23)) | ~ spl30_581), inference(forward_demodulation, [], [f1310, f5314])).
fof(f5314, plain, ((e22 = h4(e13)) | ~ spl30_581), inference(avatar_component_clause, [], [f5313])).
fof(f1310, plain, (op2(e23, e23) = h4(e13)), inference(cnf_transformation, [], [f29])).
fof(f29, plain, ((h4(e15) = op2(op2(op2(op2(e23, e23), e23), e23), e23)) & (op2(e23, e23) = h4(e13)) & (op2(op2(op2(e23, e23), e23), e23) = h4(e12)) & (op2(op2(e23, e23), e23) = h4(e11)) & (h4(e10) = op2(op2(op2(op2(e23, e23), e23), e23), op2(e23, e23))) & (e23 = h4(e14))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG026+1.p', ax29)).
fof(f17150, plain, (spl30_169 | ~ spl30_80 | ~ spl30_121 | ~ spl30_153), inference(avatar_split_clause, [], [f17149, f2174, f2041, f1870, f2241])).
fof(f17149, plain, ((e10 = op1(e11, e11)) | (~ spl30_80 | ~ spl30_121 | ~ spl30_153)), inference(forward_demodulation, [], [f12006, f2043])).
fof(f12006, plain, ((op1(e11, e11) = op1(e12, e13)) | (~ spl30_80 | ~ spl30_153)), inference(forward_demodulation, [], [f10090, f1872])).
fof(f10090, plain, ((op1(e12, e13) = op1(op1(e13, e14), e11)) | ~ spl30_153), inference(forward_demodulation, [], [f4066, f10085])).
fof(f10085, plain, ((op1(e12, e13) = op1(e11, op1(e13, e14))) | ~ spl30_153), inference(forward_demodulation, [], [f3891, f2176])).
fof(f3891, plain, (op1(e11, op1(e13, e14)) = op1(op1(e11, e14), e13)), inference(backward_demodulation, [], [f197, f1233])).
fof(f1233, plain, (op1(e13, e14) = op1(e14, e13)), inference(cnf_transformation, [], [f22])).
fof(f22, plain, ((op1(e15, e15) = op1(e15, e15)) & (op1(e14, e15) = op1(e15, e14)) & (op1(e13, e15) = op1(e15, e13)) & (op1(e12, e15) = op1(e15, e12)) & (op1(e11, e15) = op1(e15, e11)) & (op1(e10, e15) = op1(e15, e10)) & (op1(e14, e15) = op1(e15, e14)) & (op1(e14, e14) = op1(e14, e14)) & (op1(e13, e14) = op1(e14, e13)) & (op1(e12, e14) = op1(e14, e12)) & (op1(e11, e14) = op1(e14, e11)) & (op1(e10, e14) = op1(e14, e10)) & (op1(e13, e15) = op1(e15, e13)) & (op1(e13, e14) = op1(e14, e13)) & (op1(e13, e13) = op1(e13, e13)) & (op1(e12, e13) = op1(e13, e12)) & (op1(e11, e13) = op1(e13, e11)) & (op1(e10, e13) = op1(e13, e10)) & (op1(e12, e15) = op1(e15, e12)) & (op1(e12, e14) = op1(e14, e12)) & (op1(e12, e13) = op1(e13, e12)) & (op1(e12, e12) = op1(e12, e12)) & (op1(e11, e12) = op1(e12, e11)) & (op1(e10, e12) = op1(e12, e10)) & (op1(e11, e15) = op1(e15, e11)) & (op1(e11, e14) = op1(e14, e11)) & (op1(e11, e13) = op1(e13, e11)) & (op1(e11, e12) = op1(e12, e11)) & (op1(e11, e11) = op1(e11, e11)) & (op1(e10, e11) = op1(e11, e10)) & (op1(e10, e15) = op1(e15, e10)) & (op1(e10, e14) = op1(e14, e10)) & (op1(e10, e13) = op1(e13, e10)) & (op1(e10, e12) = op1(e12, e10)) & (op1(e10, e11) = op1(e11, e10)) & (op1(e10, e10) = op1(e10, e10))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG026+1.p', ax22)).
fof(f197, plain, (op1(op1(e11, e14), e13) = op1(e11, op1(e14, e13))), inference(cnf_transformation, [], [f2])).
fof(f4066, plain, (op1(e11, op1(e13, e14)) = op1(op1(e13, e14), e11)), inference(backward_demodulation, [], [f3949, f4065])).
fof(f4065, plain, (op1(e11, op1(e13, e14)) = op1(e13, op1(e11, e14))), inference(forward_demodulation, [], [f4047, f192])).
fof(f192, plain, (op1(op1(e11, e13), e14) = op1(e11, op1(e13, e14))), inference(cnf_transformation, [], [f2])).
fof(f4047, plain, (op1(op1(e11, e13), e14) = op1(e13, op1(e11, e14))), inference(backward_demodulation, [], [f252, f1225])).
fof(f1225, plain, (op1(e11, e13) = op1(e13, e11)), inference(cnf_transformation, [], [f22])).
fof(f252, plain, (op1(op1(e13, e11), e14) = op1(e13, op1(e11, e14))), inference(cnf_transformation, [], [f2])).
fof(f3949, plain, (op1(e13, op1(e11, e14)) = op1(op1(e13, e14), e11)), inference(backward_demodulation, [], [f267, f1231])).
fof(f1231, plain, (op1(e11, e14) = op1(e14, e11)), inference(cnf_transformation, [], [f22])).
fof(f267, plain, (op1(op1(e13, e14), e11) = op1(e13, op1(e14, e11))), inference(cnf_transformation, [], [f2])).
fof(f17122, plain, (spl30_211 | ~ spl30_121 | ~ spl30_196), inference(avatar_split_clause, [], [f17121, f2353, f2041, f2416])).
fof(f17121, plain, ((e10 = op1(e10, e10)) | (~ spl30_121 | ~ spl30_196)), inference(forward_demodulation, [], [f17120, f2043])).
fof(f17120, plain, ((op1(e12, e13) = op1(e10, op1(e12, e13))) | ~ spl30_196), inference(forward_demodulation, [], [f4164, f2355])).
fof(f4164, plain, (op1(e10, op1(e12, e13)) = op1(e12, op1(e10, e13))), inference(forward_demodulation, [], [f4144, f149])).
fof(f149, plain, (op1(op1(e10, e12), e13) = op1(e10, op1(e12, e13))), inference(cnf_transformation, [], [f2])).
fof(f4144, plain, (op1(op1(e10, e12), e13) = op1(e12, op1(e10, e13))), inference(backward_demodulation, [], [f209, f1218])).
fof(f1218, plain, (op1(e10, e12) = op1(e12, e10)), inference(cnf_transformation, [], [f22])).
fof(f209, plain, (op1(op1(e12, e10), e13) = op1(e12, op1(e10, e13))), inference(cnf_transformation, [], [f2])).
fof(f17042, plain, (spl30_206 | ~ spl30_80 | ~ spl30_191), inference(avatar_split_clause, [], [f17041, f2332, f1870, f2395])).
fof(f17041, plain, ((e11 = op1(e10, e11)) | (~ spl30_80 | ~ spl30_191)), inference(forward_demodulation, [], [f17040, f1872])).
fof(f17040, plain, ((op1(e10, e11) = op1(e13, e14)) | (~ spl30_80 | ~ spl30_191)), inference(forward_demodulation, [], [f12127, f2334])).
fof(f12127, plain, ((op1(e10, e11) = op1(e13, op1(e10, e14))) | ~ spl30_80), inference(forward_demodulation, [], [f4099, f1872])).
fof(f4099, plain, (op1(e10, op1(e13, e14)) = op1(e13, op1(e10, e14))), inference(forward_demodulation, [], [f4079, f156])).
fof(f156, plain, (op1(op1(e10, e13), e14) = op1(e10, op1(e13, e14))), inference(cnf_transformation, [], [f2])).
fof(f4079, plain, (op1(op1(e10, e13), e14) = op1(e13, op1(e10, e14))), inference(backward_demodulation, [], [f246, f1224])).
fof(f1224, plain, (op1(e10, e13) = op1(e13, e10)), inference(cnf_transformation, [], [f22])).
fof(f246, plain, (op1(op1(e13, e10), e14) = op1(e13, op1(e10, e14))), inference(cnf_transformation, [], [f2])).
fof(f17003, plain, (spl30_201 | ~ spl30_153 | ~ spl30_191), inference(avatar_split_clause, [], [f17002, f2332, f2174, f2374])).
fof(f2374, plain, (spl30_201 <=> (e12 = op1(e10, e12))), introduced(avatar_definition, [new_symbols(naming, [spl30_201])])).
fof(f17002, plain, ((e12 = op1(e10, e12)) | (~ spl30_153 | ~ spl30_191)), inference(forward_demodulation, [], [f17001, f2176])).
fof(f17001, plain, ((op1(e10, e12) = op1(e11, e14)) | (~ spl30_153 | ~ spl30_191)), inference(forward_demodulation, [], [f10229, f2334])).
fof(f10229, plain, ((op1(e10, e12) = op1(e11, op1(e10, e14))) | ~ spl30_153), inference(forward_demodulation, [], [f4201, f2176])).
fof(f4201, plain, (op1(e10, op1(e11, e14)) = op1(e11, op1(e10, e14))), inference(forward_demodulation, [], [f4179, f144])).
fof(f144, plain, (op1(op1(e10, e11), e14) = op1(e10, op1(e11, e14))), inference(cnf_transformation, [], [f2])).
fof(f4179, plain, (op1(op1(e10, e11), e14) = op1(e11, op1(e10, e14))), inference(backward_demodulation, [], [f174, f1212])).
fof(f1212, plain, (op1(e10, e11) = op1(e11, e10)), inference(cnf_transformation, [], [f22])).
fof(f174, plain, (op1(op1(e11, e10), e14) = op1(e11, op1(e10, e14))), inference(cnf_transformation, [], [f2])).
fof(f16876, plain, (~ spl30_201 | ~ spl30_459 | ~ spl30_558 | spl30_752), inference(avatar_contradiction_clause, [], [f16875])).
fof(f16875, plain, ($false | (~ spl30_201 | ~ spl30_459 | ~ spl30_558 | spl30_752)), inference(subsumption_resolution, [], [f16874, f3451])).
fof(f3451, plain, ((e22 = op2(e20, e22)) | ~ spl30_459), inference(avatar_component_clause, [], [f3449])).
fof(f3449, plain, (spl30_459 <=> (e22 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl30_459])])).
fof(f16874, plain, (~ (e22 = op2(e20, e22)) | (~ spl30_201 | ~ spl30_558 | spl30_752)), inference(forward_demodulation, [], [f16873, f5190])).
fof(f16873, plain, (~ (op2(e20, e22) = h5(e12)) | (~ spl30_201 | ~ spl30_558 | spl30_752)), inference(forward_demodulation, [], [f16872, f2376])).
fof(f2376, plain, ((e12 = op1(e10, e12)) | ~ spl30_201), inference(avatar_component_clause, [], [f2374])).
fof(f16872, plain, (~ (op2(e20, e22) = h5(op1(e10, e12))) | (~ spl30_558 | spl30_752)), inference(forward_demodulation, [], [f6387, f5190])).
fof(f6387, plain, (~ (h5(op1(e10, e12)) = op2(e20, h5(e12))) | spl30_752), inference(avatar_component_clause, [], [f6385])).
fof(f6385, plain, (spl30_752 <=> (h5(op1(e10, e12)) = op2(e20, h5(e12)))), introduced(avatar_definition, [new_symbols(naming, [spl30_752])])).
fof(f16706, plain, (spl30_406 | ~ spl30_378 | ~ spl30_444 | ~ spl30_558 | ~ spl30_605), inference(avatar_split_clause, [], [f16693, f5433, f5189, f3386, f3111, f3228])).
fof(f5433, plain, (spl30_605 <=> (e23 = h3(e13))), introduced(avatar_definition, [new_symbols(naming, [spl30_605])])).
fof(f16693, plain, ((e23 = op2(e21, e25)) | (~ spl30_378 | ~ spl30_444 | ~ spl30_558 | ~ spl30_605)), inference(backward_demodulation, [], [f15759, f5434])).
fof(f5434, plain, ((e23 = h3(e13)) | ~ spl30_605), inference(avatar_component_clause, [], [f5433])).
fof(f15759, plain, ((op2(e21, e25) = h3(e13)) | (~ spl30_378 | ~ spl30_444 | ~ spl30_558)), inference(forward_demodulation, [], [f1273, f15758])).
fof(f15758, plain, ((op2(e25, e21) = h3(e13)) | (~ spl30_378 | ~ spl30_444 | ~ spl30_558)), inference(forward_demodulation, [], [f15090, f15092])).
fof(f15092, plain, ((h3(e13) = op2(e20, h3(e13))) | (~ spl30_378 | ~ spl30_444 | ~ spl30_558)), inference(forward_demodulation, [], [f12555, f12370])).
fof(f12370, plain, ((op2(e21, e25) = h3(e13)) | (~ spl30_378 | ~ spl30_558)), inference(backward_demodulation, [], [f10628, f1304])).
fof(f1304, plain, (op2(e22, e22) = h3(e13)), inference(cnf_transformation, [], [f28])).
fof(f28, plain, ((h3(e15) = op2(op2(op2(op2(e22, e22), e22), e22), e22)) & (op2(e22, e22) = h3(e13)) & (op2(op2(op2(e22, e22), e22), e22) = h3(e12)) & (op2(op2(e22, e22), e22) = h3(e11)) & (h3(e10) = op2(op2(op2(op2(e22, e22), e22), e22), op2(e22, e22))) & (e22 = h3(e14))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG026+1.p', ax28)).
fof(f10628, plain, ((op2(e21, e25) = op2(e22, e22)) | (~ spl30_378 | ~ spl30_558)), inference(backward_demodulation, [], [f10393, f5190])).
fof(f10393, plain, ((op2(e21, e25) = op2(h5(e12), e22)) | ~ spl30_378), inference(forward_demodulation, [], [f4934, f10392])).
fof(f10392, plain, ((op2(e21, e25) = op2(e22, h5(e12))) | ~ spl30_378), inference(forward_demodulation, [], [f4906, f3113])).
fof(f4906, plain, (op2(e21, op2(e22, e24)) = op2(e22, h5(e12))), inference(backward_demodulation, [], [f4591, f4888])).
fof(f4591, plain, (op2(e21, op2(e22, e24)) = op2(e22, op2(e21, e24))), inference(forward_demodulation, [], [f4572, f469])).
fof(f469, plain, (op2(op2(e21, e22), e24) = op2(e21, op2(e22, e24))), inference(cnf_transformation, [], [f6])).
fof(f4572, plain, (op2(op2(e21, e22), e24) = op2(e22, op2(e21, e24))), inference(backward_demodulation, [], [f499, f1255])).
fof(f1255, plain, (op2(e21, e22) = op2(e22, e21)), inference(cnf_transformation, [], [f23])).
fof(f23, plain, ((op2(e25, e25) = op2(e25, e25)) & (op2(e24, e25) = op2(e25, e24)) & (op2(e23, e25) = op2(e25, e23)) & (op2(e22, e25) = op2(e25, e22)) & (op2(e21, e25) = op2(e25, e21)) & (op2(e20, e25) = op2(e25, e20)) & (op2(e24, e25) = op2(e25, e24)) & (op2(e24, e24) = op2(e24, e24)) & (op2(e23, e24) = op2(e24, e23)) & (op2(e22, e24) = op2(e24, e22)) & (op2(e21, e24) = op2(e24, e21)) & (op2(e20, e24) = op2(e24, e20)) & (op2(e23, e25) = op2(e25, e23)) & (op2(e23, e24) = op2(e24, e23)) & (op2(e23, e23) = op2(e23, e23)) & (op2(e22, e23) = op2(e23, e22)) & (op2(e21, e23) = op2(e23, e21)) & (op2(e20, e23) = op2(e23, e20)) & (op2(e22, e25) = op2(e25, e22)) & (op2(e22, e24) = op2(e24, e22)) & (op2(e22, e23) = op2(e23, e22)) & (op2(e22, e22) = op2(e22, e22)) & (op2(e21, e22) = op2(e22, e21)) & (op2(e20, e22) = op2(e22, e20)) & (op2(e21, e25) = op2(e25, e21)) & (op2(e21, e24) = op2(e24, e21)) & (op2(e21, e23) = op2(e23, e21)) & (op2(e21, e22) = op2(e22, e21)) & (op2(e21, e21) = op2(e21, e21)) & (op2(e20, e21) = op2(e21, e20)) & (op2(e20, e25) = op2(e25, e20)) & (op2(e20, e24) = op2(e24, e20)) & (op2(e20, e23) = op2(e23, e20)) & (op2(e20, e22) = op2(e22, e20)) & (op2(e20, e21) = op2(e21, e20)) & (op2(e20, e20) = op2(e20, e20))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG026+1.p', ax23)).
fof(f499, plain, (op2(op2(e22, e21), e24) = op2(e22, op2(e21, e24))), inference(cnf_transformation, [], [f6])).
fof(f4934, plain, (op2(h5(e12), e22) = op2(e22, h5(e12))), inference(backward_demodulation, [], [f4901, f4906])).
fof(f4901, plain, (op2(e21, op2(e22, e24)) = op2(h5(e12), e22)), inference(backward_demodulation, [], [f4378, f4888])).
fof(f4378, plain, (op2(e21, op2(e22, e24)) = op2(op2(e21, e24), e22)), inference(backward_demodulation, [], [f479, f1268])).
fof(f1268, plain, (op2(e22, e24) = op2(e24, e22)), inference(cnf_transformation, [], [f23])).
fof(f479, plain, (op2(op2(e21, e24), e22) = op2(e21, op2(e24, e22))), inference(cnf_transformation, [], [f6])).
fof(f12555, plain, ((op2(e21, e25) = op2(e20, op2(e21, e25))) | ~ spl30_444), inference(forward_demodulation, [], [f4663, f3388])).
fof(f4663, plain, (op2(e20, op2(e21, e25)) = op2(e21, op2(e20, e25))), inference(forward_demodulation, [], [f4640, f428])).
fof(f428, plain, (op2(op2(e20, e21), e25) = op2(e20, op2(e21, e25))), inference(cnf_transformation, [], [f6])).
fof(f4640, plain, (op2(op2(e20, e21), e25) = op2(e21, op2(e20, e25))), inference(backward_demodulation, [], [f458, f1248])).
fof(f1248, plain, (op2(e20, e21) = op2(e21, e20)), inference(cnf_transformation, [], [f23])).
fof(f458, plain, (op2(op2(e21, e20), e25) = op2(e21, op2(e20, e25))), inference(cnf_transformation, [], [f6])).
fof(f15090, plain, ((op2(e25, e21) = op2(e20, h3(e13))) | (~ spl30_378 | ~ spl30_444 | ~ spl30_558)), inference(forward_demodulation, [], [f12553, f12370])).
fof(f12553, plain, ((op2(e25, e21) = op2(e20, op2(e21, e25))) | ~ spl30_444), inference(forward_demodulation, [], [f4288, f3388])).
fof(f4288, plain, (op2(e20, op2(e21, e25)) = op2(op2(e20, e25), e21)), inference(backward_demodulation, [], [f448, f1273])).
fof(f448, plain, (op2(op2(e20, e25), e21) = op2(e20, op2(e25, e21))), inference(cnf_transformation, [], [f6])).
fof(f1273, plain, (op2(e21, e25) = op2(e25, e21)), inference(cnf_transformation, [], [f23])).
fof(f16705, plain, (spl30_286 | ~ spl30_378 | ~ spl30_444 | ~ spl30_558 | ~ spl30_605), inference(avatar_split_clause, [], [f16692, f5433, f5189, f3386, f3111, f2728])).
fof(f16692, plain, ((e23 = op2(e25, e21)) | (~ spl30_378 | ~ spl30_444 | ~ spl30_558 | ~ spl30_605)), inference(backward_demodulation, [], [f15758, f5434])).
fof(f16663, plain, (spl30_605 | ~ spl30_378 | ~ spl30_379 | ~ spl30_444 | ~ spl30_454 | ~ spl30_459 | ~ spl30_558 | ~ spl30_581), inference(avatar_split_clause, [], [f16291, f5313, f5189, f3449, f3428, f3386, f3116, f3111, f5433])).
fof(f3116, plain, (spl30_379 <=> (e20 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl30_379])])).
fof(f16291, plain, ((e23 = h3(e13)) | (~ spl30_378 | ~ spl30_379 | ~ spl30_444 | ~ spl30_454 | ~ spl30_459 | ~ spl30_558 | ~ spl30_581)), inference(backward_demodulation, [], [f15875, f3430])).
fof(f15875, plain, ((op2(e20, e23) = h3(e13)) | (~ spl30_378 | ~ spl30_379 | ~ spl30_444 | ~ spl30_459 | ~ spl30_558 | ~ spl30_581)), inference(forward_demodulation, [], [f15874, f15755])).
fof(f15755, plain, ((op2(e22, e22) = h3(e13)) | (~ spl30_378 | ~ spl30_444 | ~ spl30_459 | ~ spl30_558)), inference(backward_demodulation, [], [f15750, f15092])).
fof(f15750, plain, ((op2(e22, e22) = op2(e20, h3(e13))) | ~ spl30_459), inference(forward_demodulation, [], [f4780, f3451])).
fof(f4780, plain, (op2(op2(e20, e22), e22) = op2(e20, h3(e13))), inference(backward_demodulation, [], [f431, f1304])).
fof(f431, plain, (op2(op2(e20, e22), e22) = op2(e20, op2(e22, e22))), inference(cnf_transformation, [], [f6])).
fof(f15874, plain, ((op2(e20, e23) = op2(e22, e22)) | (~ spl30_379 | ~ spl30_581)), inference(forward_demodulation, [], [f15873, f3118])).
fof(f3118, plain, ((e20 = op2(e22, e23)) | ~ spl30_379), inference(avatar_component_clause, [], [f3116])).
fof(f15873, plain, ((op2(e22, e22) = op2(op2(e22, e23), e23)) | ~ spl30_581), inference(forward_demodulation, [], [f4823, f5314])).
fof(f4823, plain, (op2(op2(e22, e23), e23) = op2(e22, h4(e13))), inference(backward_demodulation, [], [f510, f1310])).
fof(f510, plain, (op2(op2(e22, e23), e23) = op2(e22, op2(e23, e23))), inference(cnf_transformation, [], [f6])).
fof(f16661, plain, (~ spl30_605 | ~ spl30_130 | ~ spl30_553 | ~ spl30_558 | spl30_727), inference(avatar_split_clause, [], [f16391, f6285, f5189, f5162, f2078, f5433])).
fof(f2078, plain, (spl30_130 <=> (e13 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl30_130])])).
fof(f6285, plain, (spl30_727 <=> (h5(op1(e12, e12)) = op2(h5(e12), h5(e12)))), introduced(avatar_definition, [new_symbols(naming, [spl30_727])])).
fof(f16391, plain, (~ (e23 = h3(e13)) | (~ spl30_130 | ~ spl30_553 | ~ spl30_558 | spl30_727)), inference(forward_demodulation, [], [f16389, f5163])).
fof(f16389, plain, (~ (h3(e13) = h5(e13)) | (~ spl30_130 | ~ spl30_558 | spl30_727)), inference(backward_demodulation, [], [f15063, f2080])).
fof(f2080, plain, ((e13 = op1(e12, e12)) | ~ spl30_130), inference(avatar_component_clause, [], [f2078])).
fof(f15063, plain, (~ (h3(e13) = h5(op1(e12, e12))) | (~ spl30_558 | spl30_727)), inference(backward_demodulation, [], [f13645, f1304])).
fof(f13645, plain, (~ (op2(e22, e22) = h5(op1(e12, e12))) | (~ spl30_558 | spl30_727)), inference(forward_demodulation, [], [f6287, f5190])).
fof(f6287, plain, (~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | spl30_727), inference(avatar_component_clause, [], [f6285])).
fof(f16572, plain, (spl30_491 | ~ spl30_479), inference(avatar_split_clause, [], [f16571, f3532, f3582])).
fof(f3582, plain, (spl30_491 <=> (e24 = inv2(e24))), introduced(avatar_definition, [new_symbols(naming, [spl30_491])])).
fof(f3532, plain, (spl30_479 <=> (e24 = unit2)), introduced(avatar_definition, [new_symbols(naming, [spl30_479])])).
fof(f16571, plain, ((e24 = inv2(e24)) | ~ spl30_479), inference(forward_demodulation, [], [f722, f3534])).
fof(f3534, plain, ((e24 = unit2) | ~ spl30_479), inference(avatar_component_clause, [], [f3532])).
fof(f722, plain, (unit2 = inv2(unit2)), inference(cnf_transformation, [], [f13])).
fof(f13, plain, (unit2 = inv2(unit2)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG026+1.p', ax13)).
fof(f16551, plain, (spl30_427 | ~ spl30_338 | ~ spl30_379 | ~ spl30_558), inference(avatar_split_clause, [], [f15893, f5189, f3116, f2945, f3316])).
fof(f3316, plain, (spl30_427 <=> (e20 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl30_427])])).
fof(f2945, plain, (spl30_338 <=> (e21 = op2(e23, e24))), introduced(avatar_definition, [new_symbols(naming, [spl30_338])])).
fof(f15893, plain, ((e20 = op2(e21, e21)) | (~ spl30_338 | ~ spl30_379 | ~ spl30_558)), inference(forward_demodulation, [], [f15338, f3118])).
fof(f15338, plain, ((op2(e21, e21) = op2(e22, e23)) | (~ spl30_338 | ~ spl30_558)), inference(forward_demodulation, [], [f12319, f2947])).
fof(f2947, plain, ((e21 = op2(e23, e24)) | ~ spl30_338), inference(avatar_component_clause, [], [f2945])).
fof(f12319, plain, ((op2(e22, e23) = op2(op2(e23, e24), e21)) | ~ spl30_558), inference(backward_demodulation, [], [f10620, f10617])).
fof(f10617, plain, ((op2(e22, e23) = op2(e23, e22)) | ~ spl30_558), inference(backward_demodulation, [], [f4930, f5190])).
fof(f4930, plain, (op2(h5(e12), e23) = op2(e23, h5(e12))), inference(backward_demodulation, [], [f4900, f4905])).
fof(f4905, plain, (op2(e21, op2(e23, e24)) = op2(e23, h5(e12))), inference(backward_demodulation, [], [f4525, f4888])).
fof(f4525, plain, (op2(e21, op2(e23, e24)) = op2(e23, op2(e21, e24))), inference(forward_demodulation, [], [f4507, f475])).
fof(f475, plain, (op2(op2(e21, e23), e24) = op2(e21, op2(e23, e24))), inference(cnf_transformation, [], [f6])).
fof(f4507, plain, (op2(op2(e21, e23), e24) = op2(e23, op2(e21, e24))), inference(backward_demodulation, [], [f535, f1261])).
fof(f1261, plain, (op2(e21, e23) = op2(e23, e21)), inference(cnf_transformation, [], [f23])).
fof(f535, plain, (op2(op2(e23, e21), e24) = op2(e23, op2(e21, e24))), inference(cnf_transformation, [], [f6])).
fof(f4900, plain, (op2(e21, op2(e23, e24)) = op2(h5(e12), e23)), inference(backward_demodulation, [], [f4351, f4888])).
fof(f4351, plain, (op2(e21, op2(e23, e24)) = op2(op2(e21, e24), e23)), inference(backward_demodulation, [], [f480, f1269])).
fof(f1269, plain, (op2(e23, e24) = op2(e24, e23)), inference(cnf_transformation, [], [f23])).
fof(f480, plain, (op2(op2(e21, e24), e23) = op2(e21, op2(e24, e23))), inference(cnf_transformation, [], [f6])).
fof(f10620, plain, ((op2(e23, e22) = op2(op2(e23, e24), e21)) | ~ spl30_558), inference(backward_demodulation, [], [f4933, f5190])).
fof(f4933, plain, (op2(op2(e23, e24), e21) = op2(e23, h5(e12))), inference(backward_demodulation, [], [f4924, f4930])).
fof(f4924, plain, (op2(op2(e23, e24), e21) = op2(h5(e12), e23)), inference(backward_demodulation, [], [f4526, f4900])).
fof(f4526, plain, (op2(e21, op2(e23, e24)) = op2(op2(e23, e24), e21)), inference(backward_demodulation, [], [f4409, f4525])).
fof(f4409, plain, (op2(e23, op2(e21, e24)) = op2(op2(e23, e24), e21)), inference(backward_demodulation, [], [f550, f1267])).
fof(f1267, plain, (op2(e21, e24) = op2(e24, e21)), inference(cnf_transformation, [], [f23])).
fof(f550, plain, (op2(op2(e23, e24), e21) = op2(e23, op2(e24, e21))), inference(cnf_transformation, [], [f6])).
fof(f16424, plain, (spl30_295 | ~ spl30_378 | ~ spl30_379 | ~ spl30_553), inference(avatar_split_clause, [], [f15764, f5162, f3116, f3111, f2766])).
fof(f15764, plain, ((e20 = op2(e24, e25)) | (~ spl30_378 | ~ spl30_379 | ~ spl30_553)), inference(forward_demodulation, [], [f13777, f3118])).
fof(f13777, plain, ((op2(e22, e23) = op2(e24, e25)) | (~ spl30_378 | ~ spl30_553)), inference(forward_demodulation, [], [f10513, f5163])).
fof(f10513, plain, ((op2(e24, e25) = op2(e22, h5(e13))) | ~ spl30_378), inference(forward_demodulation, [], [f4876, f3113])).
fof(f4876, plain, (op2(e24, op2(e22, e24)) = op2(e22, h5(e13))), inference(backward_demodulation, [], [f4401, f1316])).
fof(f4401, plain, (op2(e22, op2(e24, e24)) = op2(e24, op2(e22, e24))), inference(forward_demodulation, [], [f4385, f517])).
fof(f517, plain, (op2(op2(e22, e24), e24) = op2(e22, op2(e24, e24))), inference(cnf_transformation, [], [f6])).
fof(f4385, plain, (op2(op2(e22, e24), e24) = op2(e24, op2(e22, e24))), inference(backward_demodulation, [], [f577, f1268])).
fof(f577, plain, (op2(op2(e24, e22), e24) = op2(e24, op2(e22, e24))), inference(cnf_transformation, [], [f6])).
fof(f16395, plain, (spl30_265 | ~ spl30_378 | ~ spl30_379 | ~ spl30_553), inference(avatar_split_clause, [], [f15747, f5162, f3116, f3111, f2641])).
fof(f15747, plain, ((e20 = op2(e25, e24)) | (~ spl30_378 | ~ spl30_379 | ~ spl30_553)), inference(forward_demodulation, [], [f13766, f3118])).
fof(f13766, plain, ((op2(e22, e23) = op2(e25, e24)) | (~ spl30_378 | ~ spl30_553)), inference(forward_demodulation, [], [f13765, f3113])).
fof(f13765, plain, ((op2(e22, e23) = op2(op2(e22, e24), e24)) | ~ spl30_553), inference(forward_demodulation, [], [f4863, f5163])).
fof(f4863, plain, (op2(op2(e22, e24), e24) = op2(e22, h5(e13))), inference(backward_demodulation, [], [f517, f1316])).
fof(f16223, plain, (spl30_439 | ~ spl30_480), inference(avatar_contradiction_clause, [], [f16222])).
fof(f16222, plain, ($false | (spl30_439 | ~ spl30_480)), inference(subsumption_resolution, [], [f16200, f3367])).
fof(f3367, plain, (~ (e20 = op2(e20, e25)) | spl30_439), inference(avatar_component_clause, [], [f3366])).
fof(f3366, plain, (spl30_439 <=> (e20 = op2(e20, e25))), introduced(avatar_definition, [new_symbols(naming, [spl30_439])])).
fof(f16200, plain, ((e20 = op2(e20, e25)) | ~ spl30_480), inference(backward_demodulation, [], [f634, f3538])).
fof(f3538, plain, ((e25 = unit2) | ~ spl30_480), inference(avatar_component_clause, [], [f3536])).
fof(f3536, plain, (spl30_480 <=> (e25 = unit2)), introduced(avatar_definition, [new_symbols(naming, [spl30_480])])).
fof(f634, plain, (e20 = op2(e20, unit2)), inference(cnf_transformation, [], [f7])).
fof(f7, plain, (((e25 = unit2) | (e24 = unit2) | (e23 = unit2) | (e22 = unit2) | (e21 = unit2) | (e20 = unit2)) & (e25 = op2(e25, unit2)) & (e25 = op2(unit2, e25)) & (e24 = op2(e24, unit2)) & (e24 = op2(unit2, e24)) & (e23 = op2(e23, unit2)) & (e23 = op2(unit2, e23)) & (e22 = op2(e22, unit2)) & (e22 = op2(unit2, e22)) & (e21 = op2(e21, unit2)) & (e21 = op2(unit2, e21)) & (e20 = op2(e20, unit2)) & (e20 = op2(unit2, e20))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG026+1.p', ax7)).
fof(f15978, plain, (spl30_425 | ~ spl30_335 | ~ spl30_338 | ~ spl30_378), inference(avatar_split_clause, [], [f15053, f3111, f2945, f2932, f3307])).
fof(f15053, plain, ((e24 = op2(e21, e22)) | (~ spl30_335 | ~ spl30_338 | ~ spl30_378)), inference(forward_demodulation, [], [f15052, f2934])).
fof(f15052, plain, ((op2(e21, e22) = op2(e23, e25)) | (~ spl30_338 | ~ spl30_378)), inference(forward_demodulation, [], [f12308, f2947])).
fof(f12308, plain, ((op2(e23, e25) = op2(op2(e23, e24), e22)) | ~ spl30_378), inference(backward_demodulation, [], [f4494, f10432])).
fof(f10432, plain, ((op2(e23, e25) = op2(e22, op2(e23, e24))) | ~ spl30_378), inference(forward_demodulation, [], [f4493, f3113])).
fof(f4493, plain, (op2(e22, op2(e23, e24)) = op2(e23, op2(e22, e24))), inference(forward_demodulation, [], [f4477, f511])).
fof(f511, plain, (op2(op2(e22, e23), e24) = op2(e22, op2(e23, e24))), inference(cnf_transformation, [], [f6])).
fof(f4477, plain, (op2(op2(e22, e23), e24) = op2(e23, op2(e22, e24))), inference(backward_demodulation, [], [f541, f1262])).
fof(f1262, plain, (op2(e22, e23) = op2(e23, e22)), inference(cnf_transformation, [], [f23])).
fof(f541, plain, (op2(op2(e23, e22), e24) = op2(e23, op2(e22, e24))), inference(cnf_transformation, [], [f6])).
fof(f4494, plain, (op2(e22, op2(e23, e24)) = op2(op2(e23, e24), e22)), inference(backward_demodulation, [], [f4380, f4493])).
fof(f4380, plain, (op2(e23, op2(e22, e24)) = op2(op2(e23, e24), e22)), inference(backward_demodulation, [], [f551, f1268])).
fof(f551, plain, (op2(op2(e23, e24), e22) = op2(e23, op2(e24, e22))), inference(cnf_transformation, [], [f6])).
fof(f15880, plain, (spl30_368 | ~ spl30_335 | ~ spl30_338 | ~ spl30_581), inference(avatar_split_clause, [], [f15379, f5313, f2945, f2932, f3070])).
fof(f15379, plain, ((e21 = op2(e22, e25)) | (~ spl30_335 | ~ spl30_338 | ~ spl30_581)), inference(backward_demodulation, [], [f15044, f5314])).
fof(f15044, plain, ((e21 = op2(h4(e13), e25)) | (~ spl30_335 | ~ spl30_338)), inference(forward_demodulation, [], [f4845, f15043])).
fof(f15043, plain, ((e21 = op2(e25, h4(e13))) | (~ spl30_335 | ~ spl30_338)), inference(forward_demodulation, [], [f15042, f2947])).
fof(f15042, plain, ((op2(e23, e24) = op2(e25, h4(e13))) | ~ spl30_335), inference(forward_demodulation, [], [f4834, f2934])).
fof(f4834, plain, (op2(e23, op2(e23, e25)) = op2(e25, h4(e13))), inference(backward_demodulation, [], [f4256, f1310])).
fof(f4256, plain, (op2(e23, op2(e23, e25)) = op2(e25, op2(e23, e23))), inference(forward_demodulation, [], [f4242, f4237])).
fof(f4237, plain, (op2(e23, op2(e23, e25)) = op2(op2(e23, e25), e23)), inference(backward_demodulation, [], [f558, f1275])).
fof(f1275, plain, (op2(e23, e25) = op2(e25, e23)), inference(cnf_transformation, [], [f23])).
fof(f558, plain, (op2(op2(e23, e25), e23) = op2(e23, op2(e25, e23))), inference(cnf_transformation, [], [f6])).
fof(f4242, plain, (op2(op2(e23, e25), e23) = op2(e25, op2(e23, e23))), inference(backward_demodulation, [], [f618, f1275])).
fof(f618, plain, (op2(op2(e25, e23), e23) = op2(e25, op2(e23, e23))), inference(cnf_transformation, [], [f6])).
fof(f4845, plain, (op2(h4(e13), e25) = op2(e25, h4(e13))), inference(backward_demodulation, [], [f4826, f4834])).
fof(f4826, plain, (op2(e23, op2(e23, e25)) = op2(h4(e13), e25)), inference(backward_demodulation, [], [f548, f1310])).
fof(f548, plain, (op2(op2(e23, e23), e25) = op2(e23, op2(e23, e25))), inference(cnf_transformation, [], [f6])).
fof(f15496, plain, (spl30_416 | ~ spl30_478), inference(avatar_contradiction_clause, [], [f15495])).
fof(f15495, plain, ($false | (spl30_416 | ~ spl30_478)), inference(subsumption_resolution, [], [f15459, f3271])).
fof(f3271, plain, (~ (e21 = op2(e21, e23)) | spl30_416), inference(avatar_component_clause, [], [f3270])).
fof(f3270, plain, (spl30_416 <=> (e21 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl30_416])])).
fof(f15459, plain, ((e21 = op2(e21, e23)) | ~ spl30_478), inference(backward_demodulation, [], [f636, f3530])).
fof(f3530, plain, ((e23 = unit2) | ~ spl30_478), inference(avatar_component_clause, [], [f3528])).
fof(f3528, plain, (spl30_478 <=> (e23 = unit2)), introduced(avatar_definition, [new_symbols(naming, [spl30_478])])).
fof(f636, plain, (e21 = op2(e21, unit2)), inference(cnf_transformation, [], [f7])).
fof(f15362, plain, (spl30_581 | ~ spl30_338 | ~ spl30_411 | ~ spl30_553), inference(avatar_split_clause, [], [f14402, f5162, f3249, f2945, f5313])).
fof(f3249, plain, (spl30_411 <=> (e22 = op2(e21, e24))), introduced(avatar_definition, [new_symbols(naming, [spl30_411])])).
fof(f14402, plain, ((e22 = h4(e13)) | (~ spl30_338 | ~ spl30_411 | ~ spl30_553)), inference(forward_demodulation, [], [f14388, f3251])).
fof(f3251, plain, ((e22 = op2(e21, e24)) | ~ spl30_411), inference(avatar_component_clause, [], [f3249])).
fof(f14388, plain, ((op2(e21, e24) = h4(e13)) | (~ spl30_338 | ~ spl30_553)), inference(backward_demodulation, [], [f12721, f2947])).
fof(f12721, plain, ((op2(op2(e23, e24), e24) = h4(e13)) | ~ spl30_553), inference(forward_demodulation, [], [f12705, f1310])).
fof(f12705, plain, ((op2(e23, e23) = op2(op2(e23, e24), e24)) | ~ spl30_553), inference(backward_demodulation, [], [f4864, f5163])).
fof(f4864, plain, (op2(op2(e23, e24), e24) = op2(e23, h5(e13))), inference(backward_demodulation, [], [f553, f1316])).
fof(f553, plain, (op2(op2(e23, e24), e24) = op2(e23, op2(e24, e24))), inference(cnf_transformation, [], [f6])).
fof(f15062, plain, (spl30_379 | ~ spl30_553 | ~ spl30_558), inference(avatar_split_clause, [], [f13762, f5189, f5162, f3116])).
fof(f13762, plain, ((e20 = op2(e22, e23)) | (~ spl30_553 | ~ spl30_558)), inference(forward_demodulation, [], [f10616, f5163])).
fof(f10616, plain, ((e20 = op2(e22, h5(e13))) | ~ spl30_558), inference(backward_demodulation, [], [f4915, f5190])).
fof(f14837, plain, (spl30_130 | ~ spl30_87 | ~ spl30_121 | ~ spl30_196), inference(avatar_split_clause, [], [f14836, f2353, f2041, f1899, f2078])).
fof(f14836, plain, ((e13 = op1(e12, e12)) | (~ spl30_87 | ~ spl30_121 | ~ spl30_196)), inference(forward_demodulation, [], [f14835, f2355])).
fof(f14835, plain, ((op1(e10, e13) = op1(e12, e12)) | (~ spl30_87 | ~ spl30_121)), inference(forward_demodulation, [], [f14834, f2043])).
fof(f14834, plain, ((op1(e12, e12) = op1(op1(e12, e13), e13)) | ~ spl30_87), inference(forward_demodulation, [], [f227, f1901])).
fof(f227, plain, (op1(op1(e12, e13), e13) = op1(e12, op1(e13, e13))), inference(cnf_transformation, [], [f2])).
fof(f14577, plain, (spl30_110 | ~ spl30_50 | ~ spl30_148 | ~ spl30_153), inference(avatar_split_clause, [], [f14576, f2174, f2153, f1745, f1995])).
fof(f14576, plain, ((e11 = op1(e12, e15)) | (~ spl30_50 | ~ spl30_148 | ~ spl30_153)), inference(forward_demodulation, [], [f14561, f1747])).
fof(f14561, plain, ((op1(e12, e15) = op1(e14, e13)) | (~ spl30_148 | ~ spl30_153)), inference(backward_demodulation, [], [f9746, f2155])).
fof(f9746, plain, ((op1(e12, e15) = op1(e14, op1(e11, e15))) | ~ spl30_153), inference(backward_demodulation, [], [f3973, f9727])).
fof(f9727, plain, ((op1(e12, e15) = op1(e11, op1(e14, e15))) | ~ spl30_153), inference(backward_demodulation, [], [f199, f2176])).
fof(f199, plain, (op1(op1(e11, e14), e15) = op1(e11, op1(e14, e15))), inference(cnf_transformation, [], [f2])).
fof(f3973, plain, (op1(e11, op1(e14, e15)) = op1(e14, op1(e11, e15))), inference(forward_demodulation, [], [f3955, f199])).
fof(f3955, plain, (op1(op1(e11, e14), e15) = op1(e14, op1(e11, e15))), inference(backward_demodulation, [], [f289, f1231])).
fof(f289, plain, (op1(op1(e14, e11), e15) = op1(e14, op1(e11, e15))), inference(cnf_transformation, [], [f2])).
fof(f14197, plain, (spl30_457 | ~ spl30_477), inference(avatar_split_clause, [], [f14164, f3524, f3441])).
fof(f3441, plain, (spl30_457 <=> (e20 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl30_457])])).
fof(f3524, plain, (spl30_477 <=> (e22 = unit2)), introduced(avatar_definition, [new_symbols(naming, [spl30_477])])).
fof(f14164, plain, ((e20 = op2(e20, e22)) | ~ spl30_477), inference(backward_demodulation, [], [f634, f3526])).
fof(f3526, plain, ((e22 = unit2) | ~ spl30_477), inference(avatar_component_clause, [], [f3524])).
fof(f13509, plain, (spl30_196 | ~ spl30_46 | ~ spl30_191), inference(avatar_split_clause, [], [f13508, f2332, f1728, f2353])).
fof(f13508, plain, ((e13 = op1(e10, e13)) | (~ spl30_46 | ~ spl30_191)), inference(forward_demodulation, [], [f13507, f1730])).
fof(f13507, plain, ((op1(e10, e13) = op1(e14, e14)) | (~ spl30_46 | ~ spl30_191)), inference(forward_demodulation, [], [f10302, f2334])).
fof(f10302, plain, ((op1(e10, e13) = op1(e14, op1(e10, e14))) | ~ spl30_46), inference(forward_demodulation, [], [f4005, f1730])).
fof(f4005, plain, (op1(e10, op1(e14, e14)) = op1(e14, op1(e10, e14))), inference(forward_demodulation, [], [f3985, f162])).
fof(f162, plain, (op1(op1(e10, e14), e14) = op1(e10, op1(e14, e14))), inference(cnf_transformation, [], [f2])).
fof(f3985, plain, (op1(op1(e10, e14), e14) = op1(e14, op1(e10, e14))), inference(backward_demodulation, [], [f282, f1230])).
fof(f1230, plain, (op1(e10, e14) = op1(e14, e10)), inference(cnf_transformation, [], [f22])).
fof(f282, plain, (op1(op1(e14, e10), e14) = op1(e14, op1(e10, e14))), inference(cnf_transformation, [], [f2])).
fof(f13396, plain, (~ spl30_146 | ~ spl30_206), inference(avatar_split_clause, [], [f13395, f2395, f2145])).
fof(f13395, plain, (~ (e11 = op1(e11, e15)) | ~ spl30_206), inference(forward_demodulation, [], [f4186, f2397])).
fof(f4186, plain, ~ (op1(e10, e11) = op1(e11, e15)), inference(backward_demodulation, [], [f889, f1212])).
fof(f889, plain, ~ (op1(e11, e10) = op1(e11, e15)), inference(cnf_transformation, [], [f17])).
fof(f17, plain, (~ (op1(e15, e14) = op1(e15, e15)) & ~ (op1(e15, e13) = op1(e15, e15)) & ~ (op1(e15, e13) = op1(e15, e14)) & ~ (op1(e15, e12) = op1(e15, e15)) & ~ (op1(e15, e12) = op1(e15, e14)) & ~ (op1(e15, e12) = op1(e15, e13)) & ~ (op1(e15, e11) = op1(e15, e15)) & ~ (op1(e15, e11) = op1(e15, e14)) & ~ (op1(e15, e11) = op1(e15, e13)) & ~ (op1(e15, e11) = op1(e15, e12)) & ~ (op1(e15, e10) = op1(e15, e15)) & ~ (op1(e15, e10) = op1(e15, e14)) & ~ (op1(e15, e10) = op1(e15, e13)) & ~ (op1(e15, e10) = op1(e15, e12)) & ~ (op1(e15, e10) = op1(e15, e11)) & ~ (op1(e14, e14) = op1(e14, e15)) & ~ (op1(e14, e13) = op1(e14, e15)) & ~ (op1(e14, e13) = op1(e14, e14)) & ~ (op1(e14, e12) = op1(e14, e15)) & ~ (op1(e14, e12) = op1(e14, e14)) & ~ (op1(e14, e12) = op1(e14, e13)) & ~ (op1(e14, e11) = op1(e14, e15)) & ~ (op1(e14, e11) = op1(e14, e14)) & ~ (op1(e14, e11) = op1(e14, e13)) & ~ (op1(e14, e11) = op1(e14, e12)) & ~ (op1(e14, e10) = op1(e14, e15)) & ~ (op1(e14, e10) = op1(e14, e14)) & ~ (op1(e14, e10) = op1(e14, e13)) & ~ (op1(e14, e10) = op1(e14, e12)) & ~ (op1(e14, e10) = op1(e14, e11)) & ~ (op1(e13, e14) = op1(e13, e15)) & ~ (op1(e13, e13) = op1(e13, e15)) & ~ (op1(e13, e13) = op1(e13, e14)) & ~ (op1(e13, e12) = op1(e13, e15)) & ~ (op1(e13, e12) = op1(e13, e14)) & ~ (op1(e13, e12) = op1(e13, e13)) & ~ (op1(e13, e11) = op1(e13, e15)) & ~ (op1(e13, e11) = op1(e13, e14)) & ~ (op1(e13, e11) = op1(e13, e13)) & ~ (op1(e13, e11) = op1(e13, e12)) & ~ (op1(e13, e10) = op1(e13, e15)) & ~ (op1(e13, e10) = op1(e13, e14)) & ~ (op1(e13, e10) = op1(e13, e13)) & ~ (op1(e13, e10) = op1(e13, e12)) & ~ (op1(e13, e10) = op1(e13, e11)) & ~ (op1(e12, e14) = op1(e12, e15)) & ~ (op1(e12, e13) = op1(e12, e15)) & ~ (op1(e12, e13) = op1(e12, e14)) & ~ (op1(e12, e12) = op1(e12, e15)) & ~ (op1(e12, e12) = op1(e12, e14)) & ~ (op1(e12, e12) = op1(e12, e13)) & ~ (op1(e12, e11) = op1(e12, e15)) & ~ (op1(e12, e11) = op1(e12, e14)) & ~ (op1(e12, e11) = op1(e12, e13)) & ~ (op1(e12, e11) = op1(e12, e12)) & ~ (op1(e12, e10) = op1(e12, e15)) & ~ (op1(e12, e10) = op1(e12, e14)) & ~ (op1(e12, e10) = op1(e12, e13)) & ~ (op1(e12, e10) = op1(e12, e12)) & ~ (op1(e12, e10) = op1(e12, e11)) & ~ (op1(e11, e14) = op1(e11, e15)) & ~ (op1(e11, e13) = op1(e11, e15)) & ~ (op1(e11, e13) = op1(e11, e14)) & ~ (op1(e11, e12) = op1(e11, e15)) & ~ (op1(e11, e12) = op1(e11, e14)) & ~ (op1(e11, e12) = op1(e11, e13)) & ~ (op1(e11, e11) = op1(e11, e15)) & ~ (op1(e11, e11) = op1(e11, e14)) & ~ (op1(e11, e11) = op1(e11, e13)) & ~ (op1(e11, e11) = op1(e11, e12)) & ~ (op1(e11, e10) = op1(e11, e15)) & ~ (op1(e11, e10) = op1(e11, e14)) & ~ (op1(e11, e10) = op1(e11, e13)) & ~ (op1(e11, e10) = op1(e11, e12)) & ~ (op1(e11, e10) = op1(e11, e11)) & ~ (op1(e10, e14) = op1(e10, e15)) & ~ (op1(e10, e13) = op1(e10, e15)) & ~ (op1(e10, e13) = op1(e10, e14)) & ~ (op1(e10, e12) = op1(e10, e15)) & ~ (op1(e10, e12) = op1(e10, e14)) & ~ (op1(e10, e12) = op1(e10, e13)) & ~ (op1(e10, e11) = op1(e10, e15)) & ~ (op1(e10, e11) = op1(e10, e14)) & ~ (op1(e10, e11) = op1(e10, e13)) & ~ (op1(e10, e11) = op1(e10, e12)) & ~ (op1(e10, e10) = op1(e10, e15)) & ~ (op1(e10, e10) = op1(e10, e14)) & ~ (op1(e10, e10) = op1(e10, e13)) & ~ (op1(e10, e10) = op1(e10, e12)) & ~ (op1(e10, e10) = op1(e10, e11)) & ~ (op1(e14, e15) = op1(e15, e15)) & ~ (op1(e13, e15) = op1(e15, e15)) & ~ (op1(e13, e15) = op1(e14, e15)) & ~ (op1(e12, e15) = op1(e15, e15)) & ~ (op1(e12, e15) = op1(e14, e15)) & ~ (op1(e12, e15) = op1(e13, e15)) & ~ (op1(e11, e15) = op1(e15, e15)) & ~ (op1(e11, e15) = op1(e14, e15)) & ~ (op1(e11, e15) = op1(e13, e15)) & ~ (op1(e11, e15) = op1(e12, e15)) & ~ (op1(e10, e15) = op1(e15, e15)) & ~ (op1(e10, e15) = op1(e14, e15)) & ~ (op1(e10, e15) = op1(e13, e15)) & ~ (op1(e10, e15) = op1(e12, e15)) & ~ (op1(e10, e15) = op1(e11, e15)) & ~ (op1(e14, e14) = op1(e15, e14)) & ~ (op1(e13, e14) = op1(e15, e14)) & ~ (op1(e13, e14) = op1(e14, e14)) & ~ (op1(e12, e14) = op1(e15, e14)) & ~ (op1(e12, e14) = op1(e14, e14)) & ~ (op1(e12, e14) = op1(e13, e14)) & ~ (op1(e11, e14) = op1(e15, e14)) & ~ (op1(e11, e14) = op1(e14, e14)) & ~ (op1(e11, e14) = op1(e13, e14)) & ~ (op1(e11, e14) = op1(e12, e14)) & ~ (op1(e10, e14) = op1(e15, e14)) & ~ (op1(e10, e14) = op1(e14, e14)) & ~ (op1(e10, e14) = op1(e13, e14)) & ~ (op1(e10, e14) = op1(e12, e14)) & ~ (op1(e10, e14) = op1(e11, e14)) & ~ (op1(e14, e13) = op1(e15, e13)) & ~ (op1(e13, e13) = op1(e15, e13)) & ~ (op1(e13, e13) = op1(e14, e13)) & ~ (op1(e12, e13) = op1(e15, e13)) & ~ (op1(e12, e13) = op1(e14, e13)) & ~ (op1(e12, e13) = op1(e13, e13)) & ~ (op1(e11, e13) = op1(e15, e13)) & ~ (op1(e11, e13) = op1(e14, e13)) & ~ (op1(e11, e13) = op1(e13, e13)) & ~ (op1(e11, e13) = op1(e12, e13)) & ~ (op1(e10, e13) = op1(e15, e13)) & ~ (op1(e10, e13) = op1(e14, e13)) & ~ (op1(e10, e13) = op1(e13, e13)) & ~ (op1(e10, e13) = op1(e12, e13)) & ~ (op1(e10, e13) = op1(e11, e13)) & ~ (op1(e14, e12) = op1(e15, e12)) & ~ (op1(e13, e12) = op1(e15, e12)) & ~ (op1(e13, e12) = op1(e14, e12)) & ~ (op1(e12, e12) = op1(e15, e12)) & ~ (op1(e12, e12) = op1(e14, e12)) & ~ (op1(e12, e12) = op1(e13, e12)) & ~ (op1(e11, e12) = op1(e15, e12)) & ~ (op1(e11, e12) = op1(e14, e12)) & ~ (op1(e11, e12) = op1(e13, e12)) & ~ (op1(e11, e12) = op1(e12, e12)) & ~ (op1(e10, e12) = op1(e15, e12)) & ~ (op1(e10, e12) = op1(e14, e12)) & ~ (op1(e10, e12) = op1(e13, e12)) & ~ (op1(e10, e12) = op1(e12, e12)) & ~ (op1(e10, e12) = op1(e11, e12)) & ~ (op1(e14, e11) = op1(e15, e11)) & ~ (op1(e13, e11) = op1(e15, e11)) & ~ (op1(e13, e11) = op1(e14, e11)) & ~ (op1(e12, e11) = op1(e15, e11)) & ~ (op1(e12, e11) = op1(e14, e11)) & ~ (op1(e12, e11) = op1(e13, e11)) & ~ (op1(e11, e11) = op1(e15, e11)) & ~ (op1(e11, e11) = op1(e14, e11)) & ~ (op1(e11, e11) = op1(e13, e11)) & ~ (op1(e11, e11) = op1(e12, e11)) & ~ (op1(e10, e11) = op1(e15, e11)) & ~ (op1(e10, e11) = op1(e14, e11)) & ~ (op1(e10, e11) = op1(e13, e11)) & ~ (op1(e10, e11) = op1(e12, e11)) & ~ (op1(e10, e11) = op1(e11, e11)) & ~ (op1(e14, e10) = op1(e15, e10)) & ~ (op1(e13, e10) = op1(e15, e10)) & ~ (op1(e13, e10) = op1(e14, e10)) & ~ (op1(e12, e10) = op1(e15, e10)) & ~ (op1(e12, e10) = op1(e14, e10)) & ~ (op1(e12, e10) = op1(e13, e10)) & ~ (op1(e11, e10) = op1(e15, e10)) & ~ (op1(e11, e10) = op1(e14, e10)) & ~ (op1(e11, e10) = op1(e13, e10)) & ~ (op1(e11, e10) = op1(e12, e10)) & ~ (op1(e10, e10) = op1(e15, e10)) & ~ (op1(e10, e10) = op1(e14, e10)) & ~ (op1(e10, e10) = op1(e13, e10)) & ~ (op1(e10, e10) = op1(e12, e10)) & ~ (op1(e10, e10) = op1(e11, e10))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG026+1.p', ax17)).
fof(f13288, plain, (spl30_3 | ~ spl30_120 | ~ spl30_234), inference(avatar_split_clause, [], [f13287, f2511, f2036, f1549])).
fof(f2511, plain, (spl30_234 <=> (e15 = inv1(e14))), introduced(avatar_definition, [new_symbols(naming, [spl30_234])])).
fof(f13287, plain, ((e12 = op1(e15, e15)) | (~ spl30_120 | ~ spl30_234)), inference(forward_demodulation, [], [f11949, f354])).
fof(f354, plain, (e12 = op1(unit1, e12)), inference(cnf_transformation, [], [f3])).
fof(f11949, plain, ((op1(e15, e15) = op1(unit1, e12)) | (~ spl30_120 | ~ spl30_234)), inference(forward_demodulation, [], [f9868, f11197])).
fof(f11197, plain, ((op1(e14, e15) = unit1) | ~ spl30_234), inference(backward_demodulation, [], [f371, f2513])).
fof(f2513, plain, ((e15 = inv1(e14)) | ~ spl30_234), inference(avatar_component_clause, [], [f2511])).
fof(f9868, plain, ((op1(e15, e15) = op1(op1(e14, e15), e12)) | ~ spl30_120), inference(backward_demodulation, [], [f3943, f9851])).
fof(f9851, plain, ((op1(e15, e15) = op1(e12, op1(e14, e15))) | ~ spl30_120), inference(backward_demodulation, [], [f235, f2038])).
fof(f235, plain, (op1(op1(e12, e14), e15) = op1(e12, op1(e14, e15))), inference(cnf_transformation, [], [f2])).
fof(f3943, plain, (op1(e12, op1(e14, e15)) = op1(op1(e14, e15), e12)), inference(backward_demodulation, [], [f3804, f3942])).
fof(f3942, plain, (op1(e12, op1(e14, e15)) = op1(e14, op1(e12, e15))), inference(forward_demodulation, [], [f3926, f235])).
fof(f3926, plain, (op1(op1(e12, e14), e15) = op1(e14, op1(e12, e15))), inference(backward_demodulation, [], [f295, f1232])).
fof(f1232, plain, (op1(e12, e14) = op1(e14, e12)), inference(cnf_transformation, [], [f22])).
fof(f295, plain, (op1(op1(e14, e12), e15) = op1(e14, op1(e12, e15))), inference(cnf_transformation, [], [f2])).
fof(f3804, plain, (op1(e14, op1(e12, e15)) = op1(op1(e14, e15), e12)), inference(backward_demodulation, [], [f310, f1238])).
fof(f1238, plain, (op1(e12, e15) = op1(e15, e12)), inference(cnf_transformation, [], [f22])).
fof(f310, plain, (op1(op1(e14, e15), e12) = op1(e14, op1(e15, e12))), inference(cnf_transformation, [], [f2])).
fof(f13160, plain, (spl30_37 | ~ spl30_191 | ~ spl30_234), inference(avatar_split_clause, [], [f13146, f2511, f2332, f1691])).
fof(f13146, plain, ((e10 = op1(e14, e15)) | (~ spl30_191 | ~ spl30_234)), inference(backward_demodulation, [], [f12183, f2334])).
fof(f12183, plain, ((e10 = op1(op1(e10, e14), e15)) | ~ spl30_234), inference(backward_demodulation, [], [f11873, f351])).
fof(f11873, plain, ((op1(op1(e10, e14), e15) = op1(e10, unit1)) | ~ spl30_234), inference(forward_demodulation, [], [f163, f11197])).
fof(f163, plain, (op1(op1(e10, e14), e15) = op1(e10, op1(e14, e15))), inference(cnf_transformation, [], [f2])).
fof(f13157, plain, (spl30_186 | ~ spl30_120 | ~ spl30_191), inference(avatar_split_clause, [], [f13156, f2332, f2036, f2311])).
fof(f13156, plain, ((e15 = op1(e10, e15)) | (~ spl30_120 | ~ spl30_191)), inference(forward_demodulation, [], [f13140, f2038])).
fof(f13140, plain, ((op1(e10, e15) = op1(e12, e14)) | (~ spl30_120 | ~ spl30_191)), inference(backward_demodulation, [], [f10166, f2334])).
fof(f10166, plain, ((op1(e10, e15) = op1(e12, op1(e10, e14))) | ~ spl30_120), inference(forward_demodulation, [], [f4166, f2038])).
fof(f4166, plain, (op1(e10, op1(e12, e14)) = op1(e12, op1(e10, e14))), inference(forward_demodulation, [], [f4145, f150])).
fof(f150, plain, (op1(op1(e10, e12), e14) = op1(e10, op1(e12, e14))), inference(cnf_transformation, [], [f2])).
fof(f4145, plain, (op1(op1(e10, e12), e14) = op1(e12, op1(e10, e14))), inference(backward_demodulation, [], [f210, f1218])).
fof(f210, plain, (op1(op1(e12, e10), e14) = op1(e12, op1(e10, e14))), inference(cnf_transformation, [], [f2])).
fof(f12864, plain, (spl30_335 | ~ spl30_295 | ~ spl30_378 | ~ spl30_449 | ~ spl30_553), inference(avatar_split_clause, [], [f12844, f5162, f3407, f3111, f2766, f2932])).
fof(f12844, plain, ((e24 = op2(e23, e25)) | (~ spl30_295 | ~ spl30_378 | ~ spl30_449 | ~ spl30_553)), inference(backward_demodulation, [], [f12729, f3409])).
fof(f12729, plain, ((op2(e20, e24) = op2(e23, e25)) | (~ spl30_295 | ~ spl30_378 | ~ spl30_553)), inference(backward_demodulation, [], [f12309, f12714])).
fof(f12714, plain, ((op2(e20, e24) = op2(e25, e23)) | (~ spl30_295 | ~ spl30_553)), inference(backward_demodulation, [], [f10511, f5163])).
fof(f10511, plain, ((op2(e20, e24) = op2(e25, h5(e13))) | ~ spl30_295), inference(forward_demodulation, [], [f4886, f2768])).
fof(f4886, plain, (op2(op2(e24, e25), e24) = op2(e25, h5(e13))), inference(backward_demodulation, [], [f4884, f4885])).
fof(f4885, plain, (op2(h5(e13), e25) = op2(e25, h5(e13))), inference(backward_demodulation, [], [f4865, f4873])).
fof(f4873, plain, (op2(e24, op2(e24, e25)) = op2(e25, h5(e13))), inference(backward_demodulation, [], [f4231, f1316])).
fof(f4231, plain, (op2(e24, op2(e24, e25)) = op2(e25, op2(e24, e24))), inference(forward_demodulation, [], [f4219, f4214])).
fof(f4214, plain, (op2(e24, op2(e24, e25)) = op2(op2(e24, e25), e24)), inference(backward_demodulation, [], [f595, f1276])).
fof(f1276, plain, (op2(e24, e25) = op2(e25, e24)), inference(cnf_transformation, [], [f23])).
fof(f595, plain, (op2(op2(e24, e25), e24) = op2(e24, op2(e25, e24))), inference(cnf_transformation, [], [f6])).
fof(f4219, plain, (op2(op2(e24, e25), e24) = op2(e25, op2(e24, e24))), inference(backward_demodulation, [], [f625, f1276])).
fof(f625, plain, (op2(op2(e25, e24), e24) = op2(e25, op2(e24, e24))), inference(cnf_transformation, [], [f6])).
fof(f4865, plain, (op2(e24, op2(e24, e25)) = op2(h5(e13), e25)), inference(backward_demodulation, [], [f590, f1316])).
fof(f590, plain, (op2(op2(e24, e24), e25) = op2(e24, op2(e24, e25))), inference(cnf_transformation, [], [f6])).
fof(f4884, plain, (op2(op2(e24, e25), e24) = op2(h5(e13), e25)), inference(backward_demodulation, [], [f4214, f4865])).
fof(f12309, plain, ((op2(e23, e25) = op2(e25, e23)) | ~ spl30_378), inference(forward_demodulation, [], [f10429, f10432])).
fof(f10429, plain, ((op2(e25, e23) = op2(e22, op2(e23, e24))) | ~ spl30_378), inference(forward_demodulation, [], [f4352, f3113])).
fof(f4352, plain, (op2(e22, op2(e23, e24)) = op2(op2(e22, e24), e23)), inference(backward_demodulation, [], [f516, f1269])).
fof(f516, plain, (op2(op2(e22, e24), e23) = op2(e22, op2(e24, e23))), inference(cnf_transformation, [], [f6])).
fof(f12863, plain, (spl30_275 | ~ spl30_295 | ~ spl30_449 | ~ spl30_553), inference(avatar_split_clause, [], [f12841, f5162, f3407, f2766, f2682])).
fof(f12841, plain, ((e24 = op2(e25, e23)) | (~ spl30_295 | ~ spl30_449 | ~ spl30_553)), inference(backward_demodulation, [], [f12714, f3409])).
fof(f12810, plain, (~ spl30_411 | ~ spl30_476), inference(avatar_contradiction_clause, [], [f12809])).
fof(f12809, plain, ($false | (~ spl30_411 | ~ spl30_476)), inference(subsumption_resolution, [], [f12808, f1165])).
fof(f1165, plain, ~ (e22 = e24), inference(cnf_transformation, [], [f20])).
fof(f20, plain, (~ (e24 = e25) & ~ (e23 = e25) & ~ (e23 = e24) & ~ (e22 = e25) & ~ (e22 = e24) & ~ (e22 = e23) & ~ (e21 = e25) & ~ (e21 = e24) & ~ (e21 = e23) & ~ (e21 = e22) & ~ (e20 = e25) & ~ (e20 = e24) & ~ (e20 = e23) & ~ (e20 = e22) & ~ (e20 = e21)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG026+1.p', ax20)).
fof(f12808, plain, ((e22 = e24) | (~ spl30_411 | ~ spl30_476)), inference(forward_demodulation, [], [f12779, f3251])).
fof(f12779, plain, ((e24 = op2(e21, e24)) | ~ spl30_476), inference(backward_demodulation, [], [f641, f3522])).
fof(f3522, plain, ((e21 = unit2) | ~ spl30_476), inference(avatar_component_clause, [], [f3520])).
fof(f3520, plain, (spl30_476 <=> (e21 = unit2)), introduced(avatar_definition, [new_symbols(naming, [spl30_476])])).
fof(f641, plain, (e24 = op2(unit2, e24)), inference(cnf_transformation, [], [f7])).
fof(f12770, plain, (~ spl30_416 | ~ spl30_553 | ~ spl30_732), inference(avatar_contradiction_clause, [], [f12769])).
fof(f12769, plain, ($false | (~ spl30_416 | ~ spl30_553 | ~ spl30_732)), inference(subsumption_resolution, [], [f12768, f1163])).
fof(f1163, plain, ~ (e21 = e25), inference(cnf_transformation, [], [f20])).
fof(f12768, plain, ((e21 = e25) | (~ spl30_416 | ~ spl30_553 | ~ spl30_732)), inference(forward_demodulation, [], [f12718, f3272])).
fof(f3272, plain, ((e21 = op2(e21, e23)) | ~ spl30_416), inference(avatar_component_clause, [], [f3270])).
fof(f12718, plain, ((e25 = op2(e21, e23)) | (~ spl30_553 | ~ spl30_732)), inference(backward_demodulation, [], [f12702, f5163])).
fof(f12702, plain, ((e25 = op2(e21, h5(e13))) | ~ spl30_732), inference(backward_demodulation, [], [f4882, f6306])).
fof(f4882, plain, (h5(e15) = op2(e21, h5(e13))), inference(backward_demodulation, [], [f4859, f1316])).
fof(f4859, plain, (op2(e21, op2(e24, e24)) = h5(e15)), inference(backward_demodulation, [], [f481, f4858])).
fof(f4858, plain, (op2(op2(e21, e24), e24) = h5(e15)), inference(forward_demodulation, [], [f1317, f1284])).
fof(f1317, plain, (op2(op2(op2(op2(e24, e24), e24), e24), e24) = h5(e15)), inference(cnf_transformation, [], [f30])).
fof(f481, plain, (op2(op2(e21, e24), e24) = op2(e21, op2(e24, e24))), inference(cnf_transformation, [], [f6])).
fof(f12724, plain, (spl30_338 | ~ spl30_553), inference(avatar_split_clause, [], [f12708, f5162, f2945])).
fof(f12708, plain, ((e21 = op2(e23, e24)) | ~ spl30_553), inference(backward_demodulation, [], [f4872, f5163])).
fof(f12528, plain, (spl30_278 | ~ spl30_295 | ~ spl30_464 | ~ spl30_558), inference(avatar_split_clause, [], [f12527, f5189, f3470, f2766, f2695])).
fof(f12527, plain, ((e21 = op2(e25, e22)) | (~ spl30_295 | ~ spl30_464 | ~ spl30_558)), inference(backward_demodulation, [], [f12419, f3472])).
fof(f12419, plain, ((op2(e20, e21) = op2(e25, e22)) | (~ spl30_295 | ~ spl30_558)), inference(forward_demodulation, [], [f12418, f2768])).
fof(f12418, plain, ((op2(e25, e22) = op2(op2(e24, e25), e21)) | ~ spl30_558), inference(forward_demodulation, [], [f4923, f5190])).
fof(f4923, plain, (op2(op2(e24, e25), e21) = op2(e25, h5(e12))), inference(backward_demodulation, [], [f4916, f4920])).
fof(f4920, plain, (op2(h5(e12), e25) = op2(e25, h5(e12))), inference(backward_demodulation, [], [f4890, f4899])).
fof(f4899, plain, (op2(e21, op2(e24, e25)) = op2(e25, h5(e12))), inference(backward_demodulation, [], [f4315, f4888])).
fof(f4315, plain, (op2(e21, op2(e24, e25)) = op2(e25, op2(e21, e24))), inference(forward_demodulation, [], [f4297, f4211])).
fof(f4211, plain, (op2(e21, op2(e24, e25)) = op2(op2(e21, e25), e24)), inference(backward_demodulation, [], [f487, f1276])).
fof(f487, plain, (op2(op2(e21, e25), e24) = op2(e21, op2(e25, e24))), inference(cnf_transformation, [], [f6])).
fof(f4297, plain, (op2(op2(e21, e25), e24) = op2(e25, op2(e21, e24))), inference(backward_demodulation, [], [f607, f1273])).
fof(f607, plain, (op2(op2(e25, e21), e24) = op2(e25, op2(e21, e24))), inference(cnf_transformation, [], [f6])).
fof(f4890, plain, (op2(e21, op2(e24, e25)) = op2(h5(e12), e25)), inference(backward_demodulation, [], [f482, f4888])).
fof(f482, plain, (op2(op2(e21, e24), e25) = op2(e21, op2(e24, e25))), inference(cnf_transformation, [], [f6])).
fof(f4916, plain, (op2(op2(e24, e25), e21) = op2(h5(e12), e25)), inference(backward_demodulation, [], [f4434, f4890])).
fof(f4434, plain, (op2(e21, op2(e24, e25)) = op2(op2(e24, e25), e21)), inference(backward_demodulation, [], [f4292, f4433])).
fof(f4433, plain, (op2(e21, op2(e24, e25)) = op2(e24, op2(e21, e25))), inference(forward_demodulation, [], [f4415, f482])).
fof(f4415, plain, (op2(op2(e21, e24), e25) = op2(e24, op2(e21, e25))), inference(backward_demodulation, [], [f572, f1267])).
fof(f572, plain, (op2(op2(e24, e21), e25) = op2(e24, op2(e21, e25))), inference(cnf_transformation, [], [f6])).
fof(f4292, plain, (op2(e24, op2(e21, e25)) = op2(op2(e24, e25), e21)), inference(backward_demodulation, [], [f592, f1273])).
fof(f592, plain, (op2(op2(e24, e25), e21) = op2(e24, op2(e25, e21))), inference(cnf_transformation, [], [f6])).
fof(f12124, plain, (~ spl30_188 | ~ spl30_80), inference(avatar_split_clause, [], [f12123, f1870, f2320])).
fof(f2320, plain, (spl30_188 <=> (e11 = op1(e10, e14))), introduced(avatar_definition, [new_symbols(naming, [spl30_188])])).
fof(f12123, plain, (~ (e11 = op1(e10, e14)) | ~ spl30_80), inference(forward_demodulation, [], [f842, f1872])).
fof(f842, plain, ~ (op1(e10, e14) = op1(e13, e14)), inference(cnf_transformation, [], [f17])).
fof(f12045, plain, (spl30_77 | ~ spl30_80 | ~ spl30_120 | ~ spl30_137), inference(avatar_split_clause, [], [f12036, f2107, f2036, f1870, f1857])).
fof(f12036, plain, ((e14 = op1(e13, e15)) | (~ spl30_80 | ~ spl30_120 | ~ spl30_137)), inference(backward_demodulation, [], [f11781, f2109])).
fof(f2109, plain, ((e14 = op1(e12, e11)) | ~ spl30_137), inference(avatar_component_clause, [], [f2107])).
fof(f11781, plain, ((op1(e12, e11) = op1(e13, e15)) | (~ spl30_80 | ~ spl30_120)), inference(forward_demodulation, [], [f10063, f1872])).
fof(f10063, plain, ((op1(e13, e15) = op1(e12, op1(e13, e14))) | ~ spl30_120), inference(forward_demodulation, [], [f4033, f2038])).
fof(f4033, plain, (op1(e12, op1(e13, e14)) = op1(e13, op1(e12, e14))), inference(forward_demodulation, [], [f4017, f228])).
fof(f4017, plain, (op1(op1(e12, e13), e14) = op1(e13, op1(e12, e14))), inference(backward_demodulation, [], [f258, f1226])).
fof(f1226, plain, (op1(e12, e13) = op1(e13, e12)), inference(cnf_transformation, [], [f22])).
fof(f258, plain, (op1(op1(e13, e12), e14) = op1(e13, op1(e12, e14))), inference(cnf_transformation, [], [f2])).
fof(f11995, plain, (spl30_87 | ~ spl30_46 | ~ spl30_80 | ~ spl30_153), inference(avatar_split_clause, [], [f11994, f2174, f1870, f1728, f1899])).
fof(f11994, plain, ((e12 = op1(e13, e13)) | (~ spl30_46 | ~ spl30_80 | ~ spl30_153)), inference(forward_demodulation, [], [f11993, f2176])).
fof(f11993, plain, ((op1(e11, e14) = op1(e13, e13)) | (~ spl30_46 | ~ spl30_80)), inference(forward_demodulation, [], [f10340, f1872])).
fof(f10340, plain, ((op1(e13, e13) = op1(op1(e13, e14), e14)) | ~ spl30_46), inference(forward_demodulation, [], [f270, f1730])).
fof(f270, plain, (op1(op1(e13, e14), e14) = op1(e13, op1(e14, e14))), inference(cnf_transformation, [], [f2])).
fof(f11978, plain, (spl30_162 | ~ spl30_46 | ~ spl30_60 | ~ spl30_153), inference(avatar_split_clause, [], [f11977, f2174, f1786, f1728, f2211])).
fof(f11977, plain, ((e15 = op1(e11, e13)) | (~ spl30_46 | ~ spl30_60 | ~ spl30_153)), inference(forward_demodulation, [], [f10307, f1788])).
fof(f10307, plain, ((op1(e11, e13) = op1(e14, e12)) | (~ spl30_46 | ~ spl30_153)), inference(forward_demodulation, [], [f10306, f1730])).
fof(f10306, plain, ((op1(e14, e12) = op1(e11, op1(e14, e14))) | ~ spl30_153), inference(forward_demodulation, [], [f3972, f2176])).
fof(f3972, plain, (op1(e11, op1(e14, e14)) = op1(e14, op1(e11, e14))), inference(forward_demodulation, [], [f3954, f198])).
fof(f198, plain, (op1(op1(e11, e14), e14) = op1(e11, op1(e14, e14))), inference(cnf_transformation, [], [f2])).
fof(f3954, plain, (op1(op1(e11, e14), e14) = op1(e14, op1(e11, e14))), inference(backward_demodulation, [], [f288, f1231])).
fof(f288, plain, (op1(op1(e14, e11), e14) = op1(e14, op1(e11, e14))), inference(cnf_transformation, [], [f2])).
fof(f11973, plain, (spl30_50 | ~ spl30_46), inference(avatar_split_clause, [], [f10295, f1728, f1745])).
fof(f10295, plain, ((e11 = op1(e14, e13)) | ~ spl30_46), inference(forward_demodulation, [], [f4675, f1730])).
fof(f4675, plain, (e11 = op1(e14, op1(e14, e14))), inference(backward_demodulation, [], [f306, f1279])).
fof(f1279, plain, (e11 = op1(op1(e14, e14), e14)), inference(cnf_transformation, [], [f24])).
fof(f24, plain, ((e15 = op1(op1(op1(op1(e14, e14), e14), e14), e14)) & (e13 = op1(e14, e14)) & (e12 = op1(op1(op1(e14, e14), e14), e14)) & (e11 = op1(op1(e14, e14), e14)) & (e10 = op1(op1(op1(op1(e14, e14), e14), e14), op1(e14, e14)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG026+1.p', ax24)).
fof(f306, plain, (op1(op1(e14, e14), e14) = op1(e14, op1(e14, e14))), inference(cnf_transformation, [], [f2])).
fof(f11245, plain, (spl30_165 | ~ spl30_218), inference(avatar_split_clause, [], [f11212, f2445, f2224])).
fof(f2224, plain, (spl30_165 <=> (e12 = op1(e11, e12))), introduced(avatar_definition, [new_symbols(naming, [spl30_165])])).
fof(f11212, plain, ((e12 = op1(e11, e12)) | ~ spl30_218), inference(backward_demodulation, [], [f354, f2447])).
fof(f2447, plain, ((e11 = unit1) | ~ spl30_218), inference(avatar_component_clause, [], [f2445])).
fof(f10639, plain, (spl30_732 | ~ spl30_378 | ~ spl30_558), inference(avatar_split_clause, [], [f10638, f5189, f3111, f6305])).
fof(f10638, plain, ((e25 = h5(e15)) | (~ spl30_378 | ~ spl30_558)), inference(forward_demodulation, [], [f10613, f3113])).
fof(f10613, plain, ((op2(e22, e24) = h5(e15)) | ~ spl30_558), inference(backward_demodulation, [], [f4912, f5190])).
fof(f4912, plain, (h5(e15) = op2(h5(e12), e24)), inference(backward_demodulation, [], [f4858, f4888])).
fof(f10583, plain, (spl30_511 | ~ spl30_475), inference(avatar_split_clause, [], [f8229, f3516, f3666])).
fof(f3666, plain, (spl30_511 <=> (e20 = inv2(e20))), introduced(avatar_definition, [new_symbols(naming, [spl30_511])])).
fof(f3516, plain, (spl30_475 <=> (e20 = unit2)), introduced(avatar_definition, [new_symbols(naming, [spl30_475])])).
fof(f8229, plain, ((e20 = inv2(e20)) | ~ spl30_475), inference(backward_demodulation, [], [f722, f3518])).
fof(f3518, plain, ((e20 = unit2) | ~ spl30_475), inference(avatar_component_clause, [], [f3516])).
fof(f10580, plain, (spl30_653 | ~ spl30_427), inference(avatar_split_clause, [], [f10579, f3316, f5675])).
fof(f10579, plain, ((e20 = h2(e13)) | ~ spl30_427), inference(forward_demodulation, [], [f1298, f3318])).
fof(f3318, plain, ((e20 = op2(e21, e21)) | ~ spl30_427), inference(avatar_component_clause, [], [f3316])).
fof(f1298, plain, (op2(e21, e21) = h2(e13)), inference(cnf_transformation, [], [f27])).
fof(f27, plain, ((h2(e15) = op2(op2(op2(op2(e21, e21), e21), e21), e21)) & (op2(e21, e21) = h2(e13)) & (op2(op2(op2(e21, e21), e21), e21) = h2(e12)) & (op2(op2(e21, e21), e21) = h2(e11)) & (h2(e10) = op2(op2(op2(op2(e21, e21), e21), e21), op2(e21, e21))) & (e21 = h2(e14))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG026+1.p', ax27)).
fof(f10529, plain, (spl30_553 | ~ spl30_304), inference(avatar_split_clause, [], [f10528, f2803, f5162])).
fof(f2803, plain, (spl30_304 <=> (e23 = op2(e24, e24))), introduced(avatar_definition, [new_symbols(naming, [spl30_304])])).
fof(f10528, plain, ((e23 = h5(e13)) | ~ spl30_304), inference(forward_demodulation, [], [f1316, f2805])).
fof(f2805, plain, ((e23 = op2(e24, e24)) | ~ spl30_304), inference(avatar_component_clause, [], [f2803])).
fof(f10510, plain, (spl30_449 | ~ spl30_475), inference(avatar_split_clause, [], [f10509, f3516, f3407])).
fof(f10509, plain, ((e24 = op2(e20, e24)) | ~ spl30_475), inference(forward_demodulation, [], [f641, f3518])).
fof(f10470, plain, (~ spl30_457 | ~ spl30_379), inference(avatar_split_clause, [], [f10469, f3116, f3441])).
fof(f10469, plain, (~ (e20 = op2(e20, e22)) | ~ spl30_379), inference(forward_demodulation, [], [f4610, f3118])).
fof(f4610, plain, ~ (op2(e20, e22) = op2(e22, e23)), inference(backward_demodulation, [], [f1082, f1254])).
fof(f1254, plain, (op2(e20, e22) = op2(e22, e20)), inference(cnf_transformation, [], [f23])).
fof(f1082, plain, ~ (op2(e22, e20) = op2(e22, e23)), inference(cnf_transformation, [], [f18])).
fof(f18, plain, (~ (op2(e25, e24) = op2(e25, e25)) & ~ (op2(e25, e23) = op2(e25, e25)) & ~ (op2(e25, e23) = op2(e25, e24)) & ~ (op2(e25, e22) = op2(e25, e25)) & ~ (op2(e25, e22) = op2(e25, e24)) & ~ (op2(e25, e22) = op2(e25, e23)) & ~ (op2(e25, e21) = op2(e25, e25)) & ~ (op2(e25, e21) = op2(e25, e24)) & ~ (op2(e25, e21) = op2(e25, e23)) & ~ (op2(e25, e21) = op2(e25, e22)) & ~ (op2(e25, e20) = op2(e25, e25)) & ~ (op2(e25, e20) = op2(e25, e24)) & ~ (op2(e25, e20) = op2(e25, e23)) & ~ (op2(e25, e20) = op2(e25, e22)) & ~ (op2(e25, e20) = op2(e25, e21)) & ~ (op2(e24, e24) = op2(e24, e25)) & ~ (op2(e24, e23) = op2(e24, e25)) & ~ (op2(e24, e23) = op2(e24, e24)) & ~ (op2(e24, e22) = op2(e24, e25)) & ~ (op2(e24, e22) = op2(e24, e24)) & ~ (op2(e24, e22) = op2(e24, e23)) & ~ (op2(e24, e21) = op2(e24, e25)) & ~ (op2(e24, e21) = op2(e24, e24)) & ~ (op2(e24, e21) = op2(e24, e23)) & ~ (op2(e24, e21) = op2(e24, e22)) & ~ (op2(e24, e20) = op2(e24, e25)) & ~ (op2(e24, e20) = op2(e24, e24)) & ~ (op2(e24, e20) = op2(e24, e23)) & ~ (op2(e24, e20) = op2(e24, e22)) & ~ (op2(e24, e20) = op2(e24, e21)) & ~ (op2(e23, e24) = op2(e23, e25)) & ~ (op2(e23, e23) = op2(e23, e25)) & ~ (op2(e23, e23) = op2(e23, e24)) & ~ (op2(e23, e22) = op2(e23, e25)) & ~ (op2(e23, e22) = op2(e23, e24)) & ~ (op2(e23, e22) = op2(e23, e23)) & ~ (op2(e23, e21) = op2(e23, e25)) & ~ (op2(e23, e21) = op2(e23, e24)) & ~ (op2(e23, e21) = op2(e23, e23)) & ~ (op2(e23, e21) = op2(e23, e22)) & ~ (op2(e23, e20) = op2(e23, e25)) & ~ (op2(e23, e20) = op2(e23, e24)) & ~ (op2(e23, e20) = op2(e23, e23)) & ~ (op2(e23, e20) = op2(e23, e22)) & ~ (op2(e23, e20) = op2(e23, e21)) & ~ (op2(e22, e24) = op2(e22, e25)) & ~ (op2(e22, e23) = op2(e22, e25)) & ~ (op2(e22, e23) = op2(e22, e24)) & ~ (op2(e22, e22) = op2(e22, e25)) & ~ (op2(e22, e22) = op2(e22, e24)) & ~ (op2(e22, e22) = op2(e22, e23)) & ~ (op2(e22, e21) = op2(e22, e25)) & ~ (op2(e22, e21) = op2(e22, e24)) & ~ (op2(e22, e21) = op2(e22, e23)) & ~ (op2(e22, e21) = op2(e22, e22)) & ~ (op2(e22, e20) = op2(e22, e25)) & ~ (op2(e22, e20) = op2(e22, e24)) & ~ (op2(e22, e20) = op2(e22, e23)) & ~ (op2(e22, e20) = op2(e22, e22)) & ~ (op2(e22, e20) = op2(e22, e21)) & ~ (op2(e21, e24) = op2(e21, e25)) & ~ (op2(e21, e23) = op2(e21, e25)) & ~ (op2(e21, e23) = op2(e21, e24)) & ~ (op2(e21, e22) = op2(e21, e25)) & ~ (op2(e21, e22) = op2(e21, e24)) & ~ (op2(e21, e22) = op2(e21, e23)) & ~ (op2(e21, e21) = op2(e21, e25)) & ~ (op2(e21, e21) = op2(e21, e24)) & ~ (op2(e21, e21) = op2(e21, e23)) & ~ (op2(e21, e21) = op2(e21, e22)) & ~ (op2(e21, e20) = op2(e21, e25)) & ~ (op2(e21, e20) = op2(e21, e24)) & ~ (op2(e21, e20) = op2(e21, e23)) & ~ (op2(e21, e20) = op2(e21, e22)) & ~ (op2(e21, e20) = op2(e21, e21)) & ~ (op2(e20, e24) = op2(e20, e25)) & ~ (op2(e20, e23) = op2(e20, e25)) & ~ (op2(e20, e23) = op2(e20, e24)) & ~ (op2(e20, e22) = op2(e20, e25)) & ~ (op2(e20, e22) = op2(e20, e24)) & ~ (op2(e20, e22) = op2(e20, e23)) & ~ (op2(e20, e21) = op2(e20, e25)) & ~ (op2(e20, e21) = op2(e20, e24)) & ~ (op2(e20, e21) = op2(e20, e23)) & ~ (op2(e20, e21) = op2(e20, e22)) & ~ (op2(e20, e20) = op2(e20, e25)) & ~ (op2(e20, e20) = op2(e20, e24)) & ~ (op2(e20, e20) = op2(e20, e23)) & ~ (op2(e20, e20) = op2(e20, e22)) & ~ (op2(e20, e20) = op2(e20, e21)) & ~ (op2(e24, e25) = op2(e25, e25)) & ~ (op2(e23, e25) = op2(e25, e25)) & ~ (op2(e23, e25) = op2(e24, e25)) & ~ (op2(e22, e25) = op2(e25, e25)) & ~ (op2(e22, e25) = op2(e24, e25)) & ~ (op2(e22, e25) = op2(e23, e25)) & ~ (op2(e21, e25) = op2(e25, e25)) & ~ (op2(e21, e25) = op2(e24, e25)) & ~ (op2(e21, e25) = op2(e23, e25)) & ~ (op2(e21, e25) = op2(e22, e25)) & ~ (op2(e20, e25) = op2(e25, e25)) & ~ (op2(e20, e25) = op2(e24, e25)) & ~ (op2(e20, e25) = op2(e23, e25)) & ~ (op2(e20, e25) = op2(e22, e25)) & ~ (op2(e20, e25) = op2(e21, e25)) & ~ (op2(e24, e24) = op2(e25, e24)) & ~ (op2(e23, e24) = op2(e25, e24)) & ~ (op2(e23, e24) = op2(e24, e24)) & ~ (op2(e22, e24) = op2(e25, e24)) & ~ (op2(e22, e24) = op2(e24, e24)) & ~ (op2(e22, e24) = op2(e23, e24)) & ~ (op2(e21, e24) = op2(e25, e24)) & ~ (op2(e21, e24) = op2(e24, e24)) & ~ (op2(e21, e24) = op2(e23, e24)) & ~ (op2(e21, e24) = op2(e22, e24)) & ~ (op2(e20, e24) = op2(e25, e24)) & ~ (op2(e20, e24) = op2(e24, e24)) & ~ (op2(e20, e24) = op2(e23, e24)) & ~ (op2(e20, e24) = op2(e22, e24)) & ~ (op2(e20, e24) = op2(e21, e24)) & ~ (op2(e24, e23) = op2(e25, e23)) & ~ (op2(e23, e23) = op2(e25, e23)) & ~ (op2(e23, e23) = op2(e24, e23)) & ~ (op2(e22, e23) = op2(e25, e23)) & ~ (op2(e22, e23) = op2(e24, e23)) & ~ (op2(e22, e23) = op2(e23, e23)) & ~ (op2(e21, e23) = op2(e25, e23)) & ~ (op2(e21, e23) = op2(e24, e23)) & ~ (op2(e21, e23) = op2(e23, e23)) & ~ (op2(e21, e23) = op2(e22, e23)) & ~ (op2(e20, e23) = op2(e25, e23)) & ~ (op2(e20, e23) = op2(e24, e23)) & ~ (op2(e20, e23) = op2(e23, e23)) & ~ (op2(e20, e23) = op2(e22, e23)) & ~ (op2(e20, e23) = op2(e21, e23)) & ~ (op2(e24, e22) = op2(e25, e22)) & ~ (op2(e23, e22) = op2(e25, e22)) & ~ (op2(e23, e22) = op2(e24, e22)) & ~ (op2(e22, e22) = op2(e25, e22)) & ~ (op2(e22, e22) = op2(e24, e22)) & ~ (op2(e22, e22) = op2(e23, e22)) & ~ (op2(e21, e22) = op2(e25, e22)) & ~ (op2(e21, e22) = op2(e24, e22)) & ~ (op2(e21, e22) = op2(e23, e22)) & ~ (op2(e21, e22) = op2(e22, e22)) & ~ (op2(e20, e22) = op2(e25, e22)) & ~ (op2(e20, e22) = op2(e24, e22)) & ~ (op2(e20, e22) = op2(e23, e22)) & ~ (op2(e20, e22) = op2(e22, e22)) & ~ (op2(e20, e22) = op2(e21, e22)) & ~ (op2(e24, e21) = op2(e25, e21)) & ~ (op2(e23, e21) = op2(e25, e21)) & ~ (op2(e23, e21) = op2(e24, e21)) & ~ (op2(e22, e21) = op2(e25, e21)) & ~ (op2(e22, e21) = op2(e24, e21)) & ~ (op2(e22, e21) = op2(e23, e21)) & ~ (op2(e21, e21) = op2(e25, e21)) & ~ (op2(e21, e21) = op2(e24, e21)) & ~ (op2(e21, e21) = op2(e23, e21)) & ~ (op2(e21, e21) = op2(e22, e21)) & ~ (op2(e20, e21) = op2(e25, e21)) & ~ (op2(e20, e21) = op2(e24, e21)) & ~ (op2(e20, e21) = op2(e23, e21)) & ~ (op2(e20, e21) = op2(e22, e21)) & ~ (op2(e20, e21) = op2(e21, e21)) & ~ (op2(e24, e20) = op2(e25, e20)) & ~ (op2(e23, e20) = op2(e25, e20)) & ~ (op2(e23, e20) = op2(e24, e20)) & ~ (op2(e22, e20) = op2(e25, e20)) & ~ (op2(e22, e20) = op2(e24, e20)) & ~ (op2(e22, e20) = op2(e23, e20)) & ~ (op2(e21, e20) = op2(e25, e20)) & ~ (op2(e21, e20) = op2(e24, e20)) & ~ (op2(e21, e20) = op2(e23, e20)) & ~ (op2(e21, e20) = op2(e22, e20)) & ~ (op2(e20, e20) = op2(e25, e20)) & ~ (op2(e20, e20) = op2(e24, e20)) & ~ (op2(e20, e20) = op2(e23, e20)) & ~ (op2(e20, e20) = op2(e22, e20)) & ~ (op2(e20, e20) = op2(e21, e20))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG026+1.p', ax18)).
fof(f10427, plain, (~ spl30_439 | ~ spl30_295), inference(avatar_split_clause, [], [f10426, f2766, f3366])).
fof(f10426, plain, (~ (e20 = op2(e20, e25)) | ~ spl30_295), inference(forward_demodulation, [], [f1038, f2768])).
fof(f1038, plain, ~ (op2(e20, e25) = op2(e24, e25)), inference(cnf_transformation, [], [f18])).
fof(f10335, plain, (~ spl30_190 | ~ spl30_46), inference(avatar_split_clause, [], [f10334, f1728, f2328])).
fof(f2328, plain, (spl30_190 <=> (e13 = op1(e10, e14))), introduced(avatar_definition, [new_symbols(naming, [spl30_190])])).
fof(f10334, plain, (~ (e13 = op1(e10, e14)) | ~ spl30_46), inference(forward_demodulation, [], [f843, f1730])).
fof(f843, plain, ~ (op1(e10, e14) = op1(e14, e14)), inference(cnf_transformation, [], [f17])).
fof(f10323, plain, (spl30_80 | ~ spl30_46), inference(avatar_split_clause, [], [f10322, f1728, f1870])).
fof(f10322, plain, ((e11 = op1(e13, e14)) | ~ spl30_46), inference(forward_demodulation, [], [f1279, f1730])).
fof(f10294, plain, (spl30_121 | ~ spl30_46 | ~ spl30_153), inference(avatar_split_clause, [], [f10293, f2174, f1728, f2041])).
fof(f10293, plain, ((e10 = op1(e12, e13)) | (~ spl30_46 | ~ spl30_153)), inference(forward_demodulation, [], [f10292, f2176])).
fof(f10292, plain, ((e10 = op1(op1(e11, e14), e13)) | ~ spl30_46), inference(forward_demodulation, [], [f4678, f1730])).
fof(f4678, plain, (e10 = op1(op1(e11, e14), op1(e14, e14))), inference(forward_demodulation, [], [f1278, f1279])).
fof(f1278, plain, (e10 = op1(op1(op1(op1(e14, e14), e14), e14), op1(e14, e14))), inference(cnf_transformation, [], [f24])).
fof(f10153, plain, (~ spl30_189 | ~ spl30_153), inference(avatar_split_clause, [], [f10152, f2174, f2324])).
fof(f2324, plain, (spl30_189 <=> (e12 = op1(e10, e14))), introduced(avatar_definition, [new_symbols(naming, [spl30_189])])).
fof(f10152, plain, (~ (e12 = op1(e10, e14)) | ~ spl30_153), inference(forward_demodulation, [], [f840, f2176])).
fof(f840, plain, ~ (op1(e10, e14) = op1(e11, e14)), inference(cnf_transformation, [], [f17])).
fof(f10151, plain, (~ spl30_192 | ~ spl30_120), inference(avatar_split_clause, [], [f10150, f2036, f2336])).
fof(f2336, plain, (spl30_192 <=> (e15 = op1(e10, e14))), introduced(avatar_definition, [new_symbols(naming, [spl30_192])])).
fof(f10150, plain, (~ (e15 = op1(e10, e14)) | ~ spl30_120), inference(forward_demodulation, [], [f841, f2038])).
fof(f841, plain, ~ (op1(e10, e14) = op1(e12, e14)), inference(cnf_transformation, [], [f17])).
fof(f10124, plain, (~ spl30_165 | ~ spl30_153), inference(avatar_split_clause, [], [f10123, f2174, f2224])).
fof(f10123, plain, (~ (e12 = op1(e11, e12)) | ~ spl30_153), inference(forward_demodulation, [], [f895, f2176])).
fof(f895, plain, ~ (op1(e11, e12) = op1(e11, e14)), inference(cnf_transformation, [], [f17])).
fof(f10074, plain, (~ spl30_121 | ~ spl30_13), inference(avatar_split_clause, [], [f10073, f1591, f2041])).
fof(f1591, plain, (spl30_13 <=> (e10 = op1(e15, e13))), introduced(avatar_definition, [new_symbols(naming, [spl30_13])])).
fof(f10073, plain, (~ (e10 = op1(e12, e13)) | ~ spl30_13), inference(forward_demodulation, [], [f836, f1593])).
fof(f1593, plain, ((e10 = op1(e15, e13)) | ~ spl30_13), inference(avatar_component_clause, [], [f1591])).
fof(f836, plain, ~ (op1(e12, e13) = op1(e15, e13)), inference(cnf_transformation, [], [f17])).
fof(f9883, plain, (spl30_13 | ~ spl30_120 | ~ spl30_121 | ~ spl30_187), inference(avatar_split_clause, [], [f9867, f2316, f2041, f2036, f1591])).
fof(f2316, plain, (spl30_187 <=> (e10 = op1(e10, e14))), introduced(avatar_definition, [new_symbols(naming, [spl30_187])])).
fof(f9867, plain, ((e10 = op1(e15, e13)) | (~ spl30_120 | ~ spl30_121 | ~ spl30_187)), inference(backward_demodulation, [], [f9825, f2038])).
fof(f9825, plain, ((e10 = op1(op1(e12, e14), e13)) | (~ spl30_121 | ~ spl30_187)), inference(backward_demodulation, [], [f3892, f9822])).
fof(f9822, plain, ((e10 = op1(e12, op1(e13, e14))) | (~ spl30_121 | ~ spl30_187)), inference(forward_demodulation, [], [f9800, f2318])).
fof(f2318, plain, ((e10 = op1(e10, e14)) | ~ spl30_187), inference(avatar_component_clause, [], [f2316])).
fof(f9800, plain, ((op1(e10, e14) = op1(e12, op1(e13, e14))) | ~ spl30_121), inference(backward_demodulation, [], [f228, f2043])).
fof(f3892, plain, (op1(e12, op1(e13, e14)) = op1(op1(e12, e14), e13)), inference(backward_demodulation, [], [f233, f1233])).
fof(f233, plain, (op1(op1(e12, e14), e13) = op1(e12, op1(e14, e13))), inference(cnf_transformation, [], [f2])).
fof(f9874, plain, (spl30_60 | ~ spl30_120), inference(avatar_split_clause, [], [f9855, f2036, f1786])).
fof(f9855, plain, ((e15 = op1(e14, e12)) | ~ spl30_120), inference(backward_demodulation, [], [f1232, f2038])).
fof(f9752, plain, (spl30_63 | ~ spl30_153), inference(avatar_split_clause, [], [f9732, f2174, f1799])).
fof(f9732, plain, ((e12 = op1(e14, e11)) | ~ spl30_153), inference(backward_demodulation, [], [f1231, f2176])).
fof(f9748, plain, (~ spl30_117 | ~ spl30_153), inference(avatar_split_clause, [], [f9728, f2174, f2024])).
fof(f9728, plain, (~ (e12 = op1(e12, e14)) | ~ spl30_153), inference(backward_demodulation, [], [f845, f2176])).
fof(f845, plain, ~ (op1(e11, e14) = op1(e12, e14)), inference(cnf_transformation, [], [f17])).
fof(f9469, plain, (~ spl30_121 | ~ spl30_199), inference(avatar_split_clause, [], [f9438, f2366, f2041])).
fof(f9438, plain, (~ (e10 = op1(e12, e13)) | ~ spl30_199), inference(backward_demodulation, [], [f4150, f2368])).
fof(f2368, plain, ((e10 = op1(e10, e12)) | ~ spl30_199), inference(avatar_component_clause, [], [f2366])).
fof(f4150, plain, ~ (op1(e10, e12) = op1(e12, e13)), inference(backward_demodulation, [], [f902, f1218])).
fof(f902, plain, ~ (op1(e12, e10) = op1(e12, e13)), inference(cnf_transformation, [], [f17])).
fof(f9214, plain, ~ spl30_233, inference(avatar_contradiction_clause, [], [f9213])).
fof(f9213, plain, ($false | ~ spl30_233), inference(subsumption_resolution, [], [f9212, f1147])).
fof(f1147, plain, ~ (e11 = e14), inference(cnf_transformation, [], [f19])).
fof(f19, plain, (~ (e14 = e15) & ~ (e13 = e15) & ~ (e13 = e14) & ~ (e12 = e15) & ~ (e12 = e14) & ~ (e12 = e13) & ~ (e11 = e15) & ~ (e11 = e14) & ~ (e11 = e13) & ~ (e11 = e12) & ~ (e10 = e15) & ~ (e10 = e14) & ~ (e10 = e13) & ~ (e10 = e12) & ~ (e10 = e11)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG026+1.p', ax19)).
fof(f9212, plain, ((e11 = e14) | ~ spl30_233), inference(backward_demodulation, [], [f359, f9192])).
fof(f9192, plain, ((e11 = op1(e14, unit1)) | ~ spl30_233), inference(backward_demodulation, [], [f4675, f9166])).
fof(f9166, plain, ((op1(e14, e14) = unit1) | ~ spl30_233), inference(backward_demodulation, [], [f371, f2509])).
fof(f2509, plain, ((e14 = inv1(e14)) | ~ spl30_233), inference(avatar_component_clause, [], [f2507])).
fof(f2507, plain, (spl30_233 <=> (e14 = inv1(e14))), introduced(avatar_definition, [new_symbols(naming, [spl30_233])])).
fof(f359, plain, (e14 = op1(e14, unit1)), inference(cnf_transformation, [], [f3])).
fof(f8775, plain, (spl30_349 | ~ spl30_379), inference(avatar_split_clause, [], [f8750, f3116, f2991])).
fof(f8750, plain, ((e20 = op2(e23, e22)) | ~ spl30_379), inference(backward_demodulation, [], [f1262, f3118])).
fof(f8703, plain, (spl30_558 | ~ spl30_411), inference(avatar_split_clause, [], [f8702, f3249, f5189])).
fof(f8702, plain, ((e22 = h5(e12)) | ~ spl30_411), inference(backward_demodulation, [], [f4888, f3251])).
fof(f8292, plain, (spl30_683 | ~ spl30_475 | ~ spl30_511), inference(avatar_split_clause, [], [f8231, f3666, f3516, f5826])).
fof(f8231, plain, ((e20 = h1(e13)) | (~ spl30_475 | ~ spl30_511)), inference(backward_demodulation, [], [f7895, f3518])).
fof(f7895, plain, ((unit2 = h1(e13)) | ~ spl30_511), inference(backward_demodulation, [], [f1292, f7882])).
fof(f7882, plain, ((op2(e20, e20) = unit2) | ~ spl30_511), inference(backward_demodulation, [], [f646, f3668])).
fof(f3668, plain, ((e20 = inv2(e20)) | ~ spl30_511), inference(avatar_component_clause, [], [f3666])).
fof(f646, plain, (unit2 = op2(e20, inv2(e20))), inference(cnf_transformation, [], [f8])).
fof(f8, plain, (((e25 = inv2(e25)) | (e24 = inv2(e25)) | (e23 = inv2(e25)) | (e22 = inv2(e25)) | (e21 = inv2(e25)) | (e20 = inv2(e25))) & ((e25 = inv2(e24)) | (e24 = inv2(e24)) | (e23 = inv2(e24)) | (e22 = inv2(e24)) | (e21 = inv2(e24)) | (e20 = inv2(e24))) & ((e25 = inv2(e23)) | (e24 = inv2(e23)) | (e23 = inv2(e23)) | (e22 = inv2(e23)) | (e21 = inv2(e23)) | (e20 = inv2(e23))) & ((e25 = inv2(e22)) | (e24 = inv2(e22)) | (e23 = inv2(e22)) | (e22 = inv2(e22)) | (e21 = inv2(e22)) | (e20 = inv2(e22))) & ((e25 = inv2(e21)) | (e24 = inv2(e21)) | (e23 = inv2(e21)) | (e22 = inv2(e21)) | (e21 = inv2(e21)) | (e20 = inv2(e21))) & ((e25 = inv2(e20)) | (e24 = inv2(e20)) | (e23 = inv2(e20)) | (e22 = inv2(e20)) | (e21 = inv2(e20)) | (e20 = inv2(e20))) & (unit2 = op2(inv2(e25), e25)) & (unit2 = op2(e25, inv2(e25))) & (unit2 = op2(inv2(e24), e24)) & (unit2 = op2(e24, inv2(e24))) & (unit2 = op2(inv2(e23), e23)) & (unit2 = op2(e23, inv2(e23))) & (unit2 = op2(inv2(e22), e22)) & (unit2 = op2(e22, inv2(e22))) & (unit2 = op2(inv2(e21), e21)) & (unit2 = op2(e21, inv2(e21))) & (unit2 = op2(inv2(e20), e20)) & (unit2 = op2(e20, inv2(e20)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG026+1.p', ax8)).
fof(f1292, plain, (op2(e20, e20) = h1(e13)), inference(cnf_transformation, [], [f26])).
fof(f26, plain, ((h1(e15) = op2(op2(op2(op2(e20, e20), e20), e20), e20)) & (op2(e20, e20) = h1(e13)) & (op2(op2(op2(e20, e20), e20), e20) = h1(e12)) & (op2(op2(e20, e20), e20) = h1(e11)) & (h1(e10) = op2(op2(op2(op2(e20, e20), e20), e20), op2(e20, e20))) & (e20 = h1(e14))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG026+1.p', ax26)).
fof(f8290, plain, (spl30_294 | ~ spl30_475), inference(avatar_split_clause, [], [f8228, f3516, f2761])).
fof(f8228, plain, ((e25 = op2(e25, e20)) | ~ spl30_475), inference(backward_demodulation, [], [f644, f3518])).
fof(f644, plain, (e25 = op2(e25, unit2)), inference(cnf_transformation, [], [f7])).
fof(f8289, plain, (spl30_444 | ~ spl30_475), inference(avatar_split_clause, [], [f8227, f3516, f3386])).
fof(f8227, plain, ((e25 = op2(e20, e25)) | ~ spl30_475), inference(backward_demodulation, [], [f643, f3518])).
fof(f643, plain, (e25 = op2(unit2, e25)), inference(cnf_transformation, [], [f7])).
fof(f8287, plain, (spl30_454 | ~ spl30_475), inference(avatar_split_clause, [], [f8225, f3516, f3428])).
fof(f8225, plain, ((e23 = op2(e20, e23)) | ~ spl30_475), inference(backward_demodulation, [], [f639, f3518])).
fof(f639, plain, (e23 = op2(unit2, e23)), inference(cnf_transformation, [], [f7])).
fof(f8286, plain, (spl30_399 | ~ spl30_475), inference(avatar_split_clause, [], [f8224, f3516, f3199])).
fof(f8224, plain, ((e22 = op2(e22, e20)) | ~ spl30_475), inference(backward_demodulation, [], [f638, f3518])).
fof(f638, plain, (e22 = op2(e22, unit2)), inference(cnf_transformation, [], [f7])).
fof(f8285, plain, (spl30_459 | ~ spl30_475), inference(avatar_split_clause, [], [f8223, f3516, f3449])).
fof(f8223, plain, ((e22 = op2(e20, e22)) | ~ spl30_475), inference(backward_demodulation, [], [f637, f3518])).
fof(f637, plain, (e22 = op2(unit2, e22)), inference(cnf_transformation, [], [f7])).
fof(f8283, plain, (spl30_464 | ~ spl30_475), inference(avatar_split_clause, [], [f8221, f3516, f3470])).
fof(f8221, plain, ((e21 = op2(e20, e21)) | ~ spl30_475), inference(backward_demodulation, [], [f635, f3518])).
fof(f635, plain, (e21 = op2(unit2, e21)), inference(cnf_transformation, [], [f7])).
fof(f8165, plain, ~ spl30_491, inference(avatar_contradiction_clause, [], [f8164])).
fof(f8164, plain, ($false | ~ spl30_491), inference(subsumption_resolution, [], [f8163, f1162])).
fof(f1162, plain, ~ (e21 = e24), inference(cnf_transformation, [], [f20])).
fof(f8163, plain, ((e21 = e24) | ~ spl30_491), inference(backward_demodulation, [], [f642, f8143])).
fof(f8143, plain, ((e21 = op2(e24, unit2)) | ~ spl30_491), inference(backward_demodulation, [], [f4880, f8126])).
fof(f8126, plain, ((unit2 = h5(e13)) | ~ spl30_491), inference(backward_demodulation, [], [f1316, f8121])).
fof(f8121, plain, ((op2(e24, e24) = unit2) | ~ spl30_491), inference(backward_demodulation, [], [f654, f3584])).
fof(f3584, plain, ((e24 = inv2(e24)) | ~ spl30_491), inference(avatar_component_clause, [], [f3582])).
fof(f654, plain, (unit2 = op2(e24, inv2(e24))), inference(cnf_transformation, [], [f8])).
fof(f4880, plain, (e21 = op2(e24, h5(e13))), inference(backward_demodulation, [], [f4684, f1316])).
fof(f4684, plain, (e21 = op2(e24, op2(e24, e24))), inference(backward_demodulation, [], [f589, f1284])).
fof(f589, plain, (op2(op2(e24, e24), e24) = op2(e24, op2(e24, e24))), inference(cnf_transformation, [], [f6])).
fof(f642, plain, (e24 = op2(e24, unit2)), inference(cnf_transformation, [], [f7])).
fof(f6396, plain, (~ spl30_727 | ~ spl30_728 | ~ spl30_729 | ~ spl30_730 | ~ spl30_731 | spl30_563 | spl30_559 | spl30_555 | spl30_551 | spl30_547 | ~ spl30_732 | ~ spl30_733 | ~ spl30_734 | ~ spl30_735 | ~ spl30_736 | ~ spl30_737 | ~ spl30_738 | ~ spl30_739 | ~ spl30_740 | ~ spl30_741 | ~ spl30_742 | ~ spl30_743 | ~ spl30_744 | ~ spl30_745 | ~ spl30_746 | ~ spl30_747 | ~ spl30_748 | ~ spl30_749 | ~ spl30_750 | ~ spl30_751 | ~ spl30_752 | ~ spl30_753 | ~ spl30_754), inference(avatar_split_clause, [], [f6283, f6393, f6389, f6385, f6381, f6377, f6373, f6369, f6365, f6361, f6357, f6353, f6349, f6345, f6341, f6337, f6333, f6329, f6325, f6321, f6317, f6313, f6309, f6305, f5129, f5152, f5174, f5196, f5219, f6301, f6297, f6293, f6289, f6285])).
fof(f5219, plain, (spl30_563 <=> sP20), introduced(avatar_definition, [new_symbols(naming, [spl30_563])])).
fof(f5196, plain, (spl30_559 <=> sP21), introduced(avatar_definition, [new_symbols(naming, [spl30_559])])).
fof(f5174, plain, (spl30_555 <=> sP22), introduced(avatar_definition, [new_symbols(naming, [spl30_555])])).
fof(f5152, plain, (spl30_551 <=> sP23), introduced(avatar_definition, [new_symbols(naming, [spl30_551])])).
fof(f5129, plain, (spl30_547 <=> sP24), introduced(avatar_definition, [new_symbols(naming, [spl30_547])])).
fof(f6283, plain, (~ (h1(e13) = h5(op1(e10, e10))) | ~ (op2(e20, e21) = h5(op1(e10, e11))) | ~ (h5(op1(e10, e12)) = op2(e20, h5(e12))) | ~ (h5(op1(e10, e13)) = op2(e20, h5(e13))) | ~ (op2(e20, e24) = h5(op1(e10, e14))) | ~ (h5(op1(e10, e15)) = op2(e20, h5(e15))) | ~ (h2(e13) = h5(op1(e11, e11))) | ~ (h5(op1(e11, e12)) = op2(e21, h5(e12))) | ~ (h5(e15) = h5(op1(e11, e13))) | ~ (h5(e12) = h5(op1(e11, e14))) | ~ (h5(op1(e11, e15)) = op2(e21, h5(e15))) | ~ (e20 = h5(op1(e12, e13))) | ~ (h5(e15) = h5(op1(e12, e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e13), h5(e12))) | ~ (e21 = h5(op1(e13, e14))) | ~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12)))), inference(forward_demodulation, [], [f6282, f1292])).
fof(f6282, plain, (~ (op2(e20, e20) = h5(op1(e10, e10))) | ~ (op2(e20, e21) = h5(op1(e10, e11))) | ~ (h5(op1(e10, e12)) = op2(e20, h5(e12))) | ~ (h5(op1(e10, e13)) = op2(e20, h5(e13))) | ~ (op2(e20, e24) = h5(op1(e10, e14))) | ~ (h5(op1(e10, e15)) = op2(e20, h5(e15))) | ~ (h2(e13) = h5(op1(e11, e11))) | ~ (h5(op1(e11, e12)) = op2(e21, h5(e12))) | ~ (h5(e15) = h5(op1(e11, e13))) | ~ (h5(e12) = h5(op1(e11, e14))) | ~ (h5(op1(e11, e15)) = op2(e21, h5(e15))) | ~ (e20 = h5(op1(e12, e13))) | ~ (h5(e15) = h5(op1(e12, e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e13), h5(e12))) | ~ (e21 = h5(op1(e13, e14))) | ~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12)))), inference(forward_demodulation, [], [f6281, f4945])).
fof(f6281, plain, (~ (op2(e20, e21) = h5(op1(e10, e11))) | ~ (h5(op1(e10, e12)) = op2(e20, h5(e12))) | ~ (h5(op1(e10, e13)) = op2(e20, h5(e13))) | ~ (op2(e20, e24) = h5(op1(e10, e14))) | ~ (h5(op1(e10, e15)) = op2(e20, h5(e15))) | ~ (h2(e13) = h5(op1(e11, e11))) | ~ (h5(op1(e11, e12)) = op2(e21, h5(e12))) | ~ (h5(e15) = h5(op1(e11, e13))) | ~ (h5(e12) = h5(op1(e11, e14))) | ~ (h5(op1(e11, e15)) = op2(e21, h5(e15))) | ~ (e20 = h5(op1(e12, e13))) | ~ (h5(e15) = h5(op1(e12, e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e13), h5(e12))) | ~ (e21 = h5(op1(e13, e14))) | ~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(duplicate_literal_removal, [], [f6280])).
fof(f6280, plain, (~ (op2(e20, e21) = h5(op1(e10, e11))) | ~ (h5(op1(e10, e12)) = op2(e20, h5(e12))) | ~ (h5(op1(e10, e13)) = op2(e20, h5(e13))) | ~ (op2(e20, e24) = h5(op1(e10, e14))) | ~ (h5(op1(e10, e15)) = op2(e20, h5(e15))) | ~ (op2(e20, e21) = h5(op1(e10, e11))) | ~ (h2(e13) = h5(op1(e11, e11))) | ~ (h5(op1(e11, e12)) = op2(e21, h5(e12))) | ~ (h5(e15) = h5(op1(e11, e13))) | ~ (h5(e12) = h5(op1(e11, e14))) | ~ (h5(op1(e11, e15)) = op2(e21, h5(e15))) | ~ (e20 = h5(op1(e12, e13))) | ~ (h5(e15) = h5(op1(e12, e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e13), h5(e12))) | ~ (e21 = h5(op1(e13, e14))) | ~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f6279, f4945])).
fof(f6279, plain, (~ (h5(op1(e10, e11)) = op2(h5(e10), e21)) | ~ (h5(op1(e10, e12)) = op2(e20, h5(e12))) | ~ (h5(op1(e10, e13)) = op2(e20, h5(e13))) | ~ (op2(e20, e24) = h5(op1(e10, e14))) | ~ (h5(op1(e10, e15)) = op2(e20, h5(e15))) | ~ (op2(e20, e21) = h5(op1(e10, e11))) | ~ (h2(e13) = h5(op1(e11, e11))) | ~ (h5(op1(e11, e12)) = op2(e21, h5(e12))) | ~ (h5(e15) = h5(op1(e11, e13))) | ~ (h5(e12) = h5(op1(e11, e14))) | ~ (h5(op1(e11, e15)) = op2(e21, h5(e15))) | ~ (e20 = h5(op1(e12, e13))) | ~ (h5(e15) = h5(op1(e12, e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e13), h5(e12))) | ~ (e21 = h5(op1(e13, e14))) | ~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f6278, f4941])).
fof(f6278, plain, (~ (h5(op1(e10, e12)) = op2(e20, h5(e12))) | ~ (h5(op1(e10, e13)) = op2(e20, h5(e13))) | ~ (op2(e20, e24) = h5(op1(e10, e14))) | ~ (h5(op1(e10, e15)) = op2(e20, h5(e15))) | ~ (op2(e20, e21) = h5(op1(e10, e11))) | ~ (h2(e13) = h5(op1(e11, e11))) | ~ (h5(op1(e11, e12)) = op2(e21, h5(e12))) | ~ (h5(e15) = h5(op1(e11, e13))) | ~ (h5(e12) = h5(op1(e11, e14))) | ~ (h5(op1(e11, e15)) = op2(e21, h5(e15))) | ~ (e20 = h5(op1(e12, e13))) | ~ (h5(e15) = h5(op1(e12, e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e13), h5(e12))) | ~ (e21 = h5(op1(e13, e14))) | ~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(duplicate_literal_removal, [], [f6277])).
fof(f6277, plain, (~ (h5(op1(e10, e12)) = op2(e20, h5(e12))) | ~ (h5(op1(e10, e13)) = op2(e20, h5(e13))) | ~ (op2(e20, e24) = h5(op1(e10, e14))) | ~ (h5(op1(e10, e15)) = op2(e20, h5(e15))) | ~ (op2(e20, e21) = h5(op1(e10, e11))) | ~ (h2(e13) = h5(op1(e11, e11))) | ~ (h5(op1(e11, e12)) = op2(e21, h5(e12))) | ~ (h5(e15) = h5(op1(e11, e13))) | ~ (h5(e12) = h5(op1(e11, e14))) | ~ (h5(op1(e11, e15)) = op2(e21, h5(e15))) | ~ (h5(op1(e10, e12)) = op2(e20, h5(e12))) | ~ (e20 = h5(op1(e12, e13))) | ~ (h5(e15) = h5(op1(e12, e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e13), h5(e12))) | ~ (e21 = h5(op1(e13, e14))) | ~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f6276, f4945])).
fof(f6276, plain, (~ (h5(op1(e10, e13)) = op2(e20, h5(e13))) | ~ (op2(e20, e24) = h5(op1(e10, e14))) | ~ (h5(op1(e10, e15)) = op2(e20, h5(e15))) | ~ (op2(e20, e21) = h5(op1(e10, e11))) | ~ (h2(e13) = h5(op1(e11, e11))) | ~ (h5(op1(e11, e12)) = op2(e21, h5(e12))) | ~ (h5(e15) = h5(op1(e11, e13))) | ~ (h5(e12) = h5(op1(e11, e14))) | ~ (h5(op1(e11, e15)) = op2(e21, h5(e15))) | ~ (h5(op1(e10, e12)) = op2(e20, h5(e12))) | ~ (e20 = h5(op1(e12, e13))) | ~ (h5(e15) = h5(op1(e12, e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e13), h5(e12))) | ~ (e21 = h5(op1(e13, e14))) | ~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(duplicate_literal_removal, [], [f6275])).
fof(f6275, plain, (~ (h5(op1(e10, e13)) = op2(e20, h5(e13))) | ~ (op2(e20, e24) = h5(op1(e10, e14))) | ~ (h5(op1(e10, e15)) = op2(e20, h5(e15))) | ~ (op2(e20, e21) = h5(op1(e10, e11))) | ~ (h2(e13) = h5(op1(e11, e11))) | ~ (h5(op1(e11, e12)) = op2(e21, h5(e12))) | ~ (h5(e15) = h5(op1(e11, e13))) | ~ (h5(e12) = h5(op1(e11, e14))) | ~ (h5(op1(e11, e15)) = op2(e21, h5(e15))) | ~ (h5(op1(e10, e12)) = op2(e20, h5(e12))) | ~ (e20 = h5(op1(e12, e13))) | ~ (h5(e15) = h5(op1(e12, e14))) | ~ (h5(op1(e10, e13)) = op2(e20, h5(e13))) | ~ (h5(op1(e12, e13)) = op2(h5(e13), h5(e12))) | ~ (e21 = h5(op1(e13, e14))) | ~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f6274, f4945])).
fof(f6274, plain, (~ (op2(e20, e24) = h5(op1(e10, e14))) | ~ (h5(op1(e10, e15)) = op2(e20, h5(e15))) | ~ (op2(e20, e21) = h5(op1(e10, e11))) | ~ (h2(e13) = h5(op1(e11, e11))) | ~ (h5(op1(e11, e12)) = op2(e21, h5(e12))) | ~ (h5(e15) = h5(op1(e11, e13))) | ~ (h5(e12) = h5(op1(e11, e14))) | ~ (h5(op1(e11, e15)) = op2(e21, h5(e15))) | ~ (h5(op1(e10, e12)) = op2(e20, h5(e12))) | ~ (e20 = h5(op1(e12, e13))) | ~ (h5(e15) = h5(op1(e12, e14))) | ~ (h5(op1(e10, e13)) = op2(e20, h5(e13))) | ~ (h5(op1(e12, e13)) = op2(h5(e13), h5(e12))) | ~ (e21 = h5(op1(e13, e14))) | ~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(duplicate_literal_removal, [], [f6273])).
fof(f6273, plain, (~ (op2(e20, e24) = h5(op1(e10, e14))) | ~ (h5(op1(e10, e15)) = op2(e20, h5(e15))) | ~ (op2(e20, e21) = h5(op1(e10, e11))) | ~ (h2(e13) = h5(op1(e11, e11))) | ~ (h5(op1(e11, e12)) = op2(e21, h5(e12))) | ~ (h5(e15) = h5(op1(e11, e13))) | ~ (h5(e12) = h5(op1(e11, e14))) | ~ (h5(op1(e11, e15)) = op2(e21, h5(e15))) | ~ (h5(op1(e10, e12)) = op2(e20, h5(e12))) | ~ (e20 = h5(op1(e12, e13))) | ~ (h5(e15) = h5(op1(e12, e14))) | ~ (h5(op1(e10, e13)) = op2(e20, h5(e13))) | ~ (h5(op1(e12, e13)) = op2(h5(e13), h5(e12))) | ~ (e21 = h5(op1(e13, e14))) | ~ (op2(e20, e24) = h5(op1(e10, e14))) | ~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f6272, f4945])).
fof(f6272, plain, (~ (h5(op1(e10, e14)) = op2(h5(e10), e24)) | ~ (h5(op1(e10, e15)) = op2(e20, h5(e15))) | ~ (op2(e20, e21) = h5(op1(e10, e11))) | ~ (h2(e13) = h5(op1(e11, e11))) | ~ (h5(op1(e11, e12)) = op2(e21, h5(e12))) | ~ (h5(e15) = h5(op1(e11, e13))) | ~ (h5(e12) = h5(op1(e11, e14))) | ~ (h5(op1(e11, e15)) = op2(e21, h5(e15))) | ~ (h5(op1(e10, e12)) = op2(e20, h5(e12))) | ~ (e20 = h5(op1(e12, e13))) | ~ (h5(e15) = h5(op1(e12, e14))) | ~ (h5(op1(e10, e13)) = op2(e20, h5(e13))) | ~ (h5(op1(e12, e13)) = op2(h5(e13), h5(e12))) | ~ (e21 = h5(op1(e13, e14))) | ~ (op2(e20, e24) = h5(op1(e10, e14))) | ~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f6271, f1312])).
fof(f6271, plain, (~ (h5(op1(e10, e15)) = op2(e20, h5(e15))) | ~ (op2(e20, e21) = h5(op1(e10, e11))) | ~ (h2(e13) = h5(op1(e11, e11))) | ~ (h5(op1(e11, e12)) = op2(e21, h5(e12))) | ~ (h5(e15) = h5(op1(e11, e13))) | ~ (h5(e12) = h5(op1(e11, e14))) | ~ (h5(op1(e11, e15)) = op2(e21, h5(e15))) | ~ (h5(op1(e10, e12)) = op2(e20, h5(e12))) | ~ (e20 = h5(op1(e12, e13))) | ~ (h5(e15) = h5(op1(e12, e14))) | ~ (h5(op1(e10, e13)) = op2(e20, h5(e13))) | ~ (h5(op1(e12, e13)) = op2(h5(e13), h5(e12))) | ~ (e21 = h5(op1(e13, e14))) | ~ (op2(e20, e24) = h5(op1(e10, e14))) | ~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f6270, f4945])).
fof(f6270, plain, (~ (op2(e20, e21) = h5(op1(e10, e11))) | ~ (h2(e13) = h5(op1(e11, e11))) | ~ (h5(op1(e11, e12)) = op2(e21, h5(e12))) | ~ (h5(e15) = h5(op1(e11, e13))) | ~ (h5(e12) = h5(op1(e11, e14))) | ~ (h5(op1(e11, e15)) = op2(e21, h5(e15))) | ~ (h5(op1(e10, e12)) = op2(e20, h5(e12))) | ~ (e20 = h5(op1(e12, e13))) | ~ (h5(e15) = h5(op1(e12, e14))) | ~ (h5(op1(e10, e13)) = op2(e20, h5(e13))) | ~ (h5(op1(e12, e13)) = op2(h5(e13), h5(e12))) | ~ (e21 = h5(op1(e13, e14))) | ~ (op2(e20, e24) = h5(op1(e10, e14))) | ~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f6269, f1248])).
fof(f6269, plain, (~ (op2(e21, e20) = h5(op1(e10, e11))) | ~ (h2(e13) = h5(op1(e11, e11))) | ~ (h5(op1(e11, e12)) = op2(e21, h5(e12))) | ~ (h5(e15) = h5(op1(e11, e13))) | ~ (h5(e12) = h5(op1(e11, e14))) | ~ (h5(op1(e11, e15)) = op2(e21, h5(e15))) | ~ (h5(op1(e10, e12)) = op2(e20, h5(e12))) | ~ (e20 = h5(op1(e12, e13))) | ~ (h5(e15) = h5(op1(e12, e14))) | ~ (h5(op1(e10, e13)) = op2(e20, h5(e13))) | ~ (h5(op1(e12, e13)) = op2(h5(e13), h5(e12))) | ~ (e21 = h5(op1(e13, e14))) | ~ (op2(e20, e24) = h5(op1(e10, e14))) | ~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f6268, f1212])).
fof(f6268, plain, (~ (op2(e21, e20) = h5(op1(e11, e10))) | ~ (h2(e13) = h5(op1(e11, e11))) | ~ (h5(op1(e11, e12)) = op2(e21, h5(e12))) | ~ (h5(e15) = h5(op1(e11, e13))) | ~ (h5(e12) = h5(op1(e11, e14))) | ~ (h5(op1(e11, e15)) = op2(e21, h5(e15))) | ~ (h5(op1(e10, e12)) = op2(e20, h5(e12))) | ~ (e20 = h5(op1(e12, e13))) | ~ (h5(e15) = h5(op1(e12, e14))) | ~ (h5(op1(e10, e13)) = op2(e20, h5(e13))) | ~ (h5(op1(e12, e13)) = op2(h5(e13), h5(e12))) | ~ (e21 = h5(op1(e13, e14))) | ~ (op2(e20, e24) = h5(op1(e10, e14))) | ~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f6267, f4941])).
fof(f6267, plain, (~ (h5(op1(e11, e10)) = op2(h5(e11), e20)) | ~ (h2(e13) = h5(op1(e11, e11))) | ~ (h5(op1(e11, e12)) = op2(e21, h5(e12))) | ~ (h5(e15) = h5(op1(e11, e13))) | ~ (h5(e12) = h5(op1(e11, e14))) | ~ (h5(op1(e11, e15)) = op2(e21, h5(e15))) | ~ (h5(op1(e10, e12)) = op2(e20, h5(e12))) | ~ (e20 = h5(op1(e12, e13))) | ~ (h5(e15) = h5(op1(e12, e14))) | ~ (h5(op1(e10, e13)) = op2(e20, h5(e13))) | ~ (h5(op1(e12, e13)) = op2(h5(e13), h5(e12))) | ~ (e21 = h5(op1(e13, e14))) | ~ (op2(e20, e24) = h5(op1(e10, e14))) | ~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f6266, f4945])).
fof(f6266, plain, (~ (h2(e13) = h5(op1(e11, e11))) | ~ (h5(op1(e11, e12)) = op2(e21, h5(e12))) | ~ (h5(e15) = h5(op1(e11, e13))) | ~ (h5(e12) = h5(op1(e11, e14))) | ~ (h5(op1(e11, e15)) = op2(e21, h5(e15))) | ~ (h5(op1(e10, e12)) = op2(e20, h5(e12))) | ~ (e20 = h5(op1(e12, e13))) | ~ (h5(e15) = h5(op1(e12, e14))) | ~ (h5(op1(e10, e13)) = op2(e20, h5(e13))) | ~ (h5(op1(e12, e13)) = op2(h5(e13), h5(e12))) | ~ (e21 = h5(op1(e13, e14))) | ~ (op2(e20, e24) = h5(op1(e10, e14))) | ~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f6265, f1298])).
fof(f6265, plain, (~ (op2(e21, e21) = h5(op1(e11, e11))) | ~ (h5(op1(e11, e12)) = op2(e21, h5(e12))) | ~ (h5(e15) = h5(op1(e11, e13))) | ~ (h5(e12) = h5(op1(e11, e14))) | ~ (h5(op1(e11, e15)) = op2(e21, h5(e15))) | ~ (h5(op1(e10, e12)) = op2(e20, h5(e12))) | ~ (e20 = h5(op1(e12, e13))) | ~ (h5(e15) = h5(op1(e12, e14))) | ~ (h5(op1(e10, e13)) = op2(e20, h5(e13))) | ~ (h5(op1(e12, e13)) = op2(h5(e13), h5(e12))) | ~ (e21 = h5(op1(e13, e14))) | ~ (op2(e20, e24) = h5(op1(e10, e14))) | ~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f6264, f4941])).
fof(f6264, plain, (~ (h5(op1(e11, e12)) = op2(e21, h5(e12))) | ~ (h5(e15) = h5(op1(e11, e13))) | ~ (h5(e12) = h5(op1(e11, e14))) | ~ (h5(op1(e11, e15)) = op2(e21, h5(e15))) | ~ (h5(op1(e10, e12)) = op2(e20, h5(e12))) | ~ (e20 = h5(op1(e12, e13))) | ~ (h5(e15) = h5(op1(e12, e14))) | ~ (h5(op1(e10, e13)) = op2(e20, h5(e13))) | ~ (h5(op1(e12, e13)) = op2(h5(e13), h5(e12))) | ~ (e21 = h5(op1(e13, e14))) | ~ (op2(e20, e24) = h5(op1(e10, e14))) | ~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(duplicate_literal_removal, [], [f6263])).
fof(f6263, plain, (~ (h5(op1(e11, e12)) = op2(e21, h5(e12))) | ~ (h5(e15) = h5(op1(e11, e13))) | ~ (h5(e12) = h5(op1(e11, e14))) | ~ (h5(op1(e11, e15)) = op2(e21, h5(e15))) | ~ (h5(op1(e10, e12)) = op2(e20, h5(e12))) | ~ (h5(op1(e11, e12)) = op2(e21, h5(e12))) | ~ (e20 = h5(op1(e12, e13))) | ~ (h5(e15) = h5(op1(e12, e14))) | ~ (h5(op1(e10, e13)) = op2(e20, h5(e13))) | ~ (h5(op1(e12, e13)) = op2(h5(e13), h5(e12))) | ~ (e21 = h5(op1(e13, e14))) | ~ (op2(e20, e24) = h5(op1(e10, e14))) | ~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f6262, f4941])).
fof(f6262, plain, (~ (h5(e15) = h5(op1(e11, e13))) | ~ (h5(e12) = h5(op1(e11, e14))) | ~ (h5(op1(e11, e15)) = op2(e21, h5(e15))) | ~ (h5(op1(e10, e12)) = op2(e20, h5(e12))) | ~ (h5(op1(e11, e12)) = op2(e21, h5(e12))) | ~ (e20 = h5(op1(e12, e13))) | ~ (h5(e15) = h5(op1(e12, e14))) | ~ (h5(op1(e10, e13)) = op2(e20, h5(e13))) | ~ (h5(op1(e12, e13)) = op2(h5(e13), h5(e12))) | ~ (e21 = h5(op1(e13, e14))) | ~ (op2(e20, e24) = h5(op1(e10, e14))) | ~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(duplicate_literal_removal, [], [f6261])).
fof(f6261, plain, (~ (h5(e15) = h5(op1(e11, e13))) | ~ (h5(e12) = h5(op1(e11, e14))) | ~ (h5(op1(e11, e15)) = op2(e21, h5(e15))) | ~ (h5(op1(e10, e12)) = op2(e20, h5(e12))) | ~ (h5(op1(e11, e12)) = op2(e21, h5(e12))) | ~ (e20 = h5(op1(e12, e13))) | ~ (h5(e15) = h5(op1(e12, e14))) | ~ (h5(op1(e10, e13)) = op2(e20, h5(e13))) | ~ (h5(e15) = h5(op1(e11, e13))) | ~ (h5(op1(e12, e13)) = op2(h5(e13), h5(e12))) | ~ (e21 = h5(op1(e13, e14))) | ~ (op2(e20, e24) = h5(op1(e10, e14))) | ~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f6260, f4882])).
fof(f6260, plain, (~ (h5(op1(e11, e13)) = op2(e21, h5(e13))) | ~ (h5(e12) = h5(op1(e11, e14))) | ~ (h5(op1(e11, e15)) = op2(e21, h5(e15))) | ~ (h5(op1(e10, e12)) = op2(e20, h5(e12))) | ~ (h5(op1(e11, e12)) = op2(e21, h5(e12))) | ~ (e20 = h5(op1(e12, e13))) | ~ (h5(e15) = h5(op1(e12, e14))) | ~ (h5(op1(e10, e13)) = op2(e20, h5(e13))) | ~ (h5(e15) = h5(op1(e11, e13))) | ~ (h5(op1(e12, e13)) = op2(h5(e13), h5(e12))) | ~ (e21 = h5(op1(e13, e14))) | ~ (op2(e20, e24) = h5(op1(e10, e14))) | ~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f6259, f4941])).
fof(f6259, plain, (~ (h5(e12) = h5(op1(e11, e14))) | ~ (h5(op1(e11, e15)) = op2(e21, h5(e15))) | ~ (h5(op1(e10, e12)) = op2(e20, h5(e12))) | ~ (h5(op1(e11, e12)) = op2(e21, h5(e12))) | ~ (e20 = h5(op1(e12, e13))) | ~ (h5(e15) = h5(op1(e12, e14))) | ~ (h5(op1(e10, e13)) = op2(e20, h5(e13))) | ~ (h5(e15) = h5(op1(e11, e13))) | ~ (h5(op1(e12, e13)) = op2(h5(e13), h5(e12))) | ~ (e21 = h5(op1(e13, e14))) | ~ (op2(e20, e24) = h5(op1(e10, e14))) | ~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(duplicate_literal_removal, [], [f6258])).
fof(f6258, plain, (~ (h5(e12) = h5(op1(e11, e14))) | ~ (h5(op1(e11, e15)) = op2(e21, h5(e15))) | ~ (h5(op1(e10, e12)) = op2(e20, h5(e12))) | ~ (h5(op1(e11, e12)) = op2(e21, h5(e12))) | ~ (e20 = h5(op1(e12, e13))) | ~ (h5(e15) = h5(op1(e12, e14))) | ~ (h5(op1(e10, e13)) = op2(e20, h5(e13))) | ~ (h5(e15) = h5(op1(e11, e13))) | ~ (h5(op1(e12, e13)) = op2(h5(e13), h5(e12))) | ~ (e21 = h5(op1(e13, e14))) | ~ (op2(e20, e24) = h5(op1(e10, e14))) | ~ (h5(e12) = h5(op1(e11, e14))) | ~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f6257, f4888])).
fof(f6257, plain, (~ (op2(e21, e24) = h5(op1(e11, e14))) | ~ (h5(op1(e11, e15)) = op2(e21, h5(e15))) | ~ (h5(op1(e10, e12)) = op2(e20, h5(e12))) | ~ (h5(op1(e11, e12)) = op2(e21, h5(e12))) | ~ (e20 = h5(op1(e12, e13))) | ~ (h5(e15) = h5(op1(e12, e14))) | ~ (h5(op1(e10, e13)) = op2(e20, h5(e13))) | ~ (h5(e15) = h5(op1(e11, e13))) | ~ (h5(op1(e12, e13)) = op2(h5(e13), h5(e12))) | ~ (e21 = h5(op1(e13, e14))) | ~ (op2(e20, e24) = h5(op1(e10, e14))) | ~ (h5(e12) = h5(op1(e11, e14))) | ~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f6256, f4941])).
fof(f6256, plain, (~ (h5(op1(e11, e14)) = op2(h5(e11), e24)) | ~ (h5(op1(e11, e15)) = op2(e21, h5(e15))) | ~ (h5(op1(e10, e12)) = op2(e20, h5(e12))) | ~ (h5(op1(e11, e12)) = op2(e21, h5(e12))) | ~ (e20 = h5(op1(e12, e13))) | ~ (h5(e15) = h5(op1(e12, e14))) | ~ (h5(op1(e10, e13)) = op2(e20, h5(e13))) | ~ (h5(e15) = h5(op1(e11, e13))) | ~ (h5(op1(e12, e13)) = op2(h5(e13), h5(e12))) | ~ (e21 = h5(op1(e13, e14))) | ~ (op2(e20, e24) = h5(op1(e10, e14))) | ~ (h5(e12) = h5(op1(e11, e14))) | ~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f6255, f1312])).
fof(f6255, plain, (~ (h5(op1(e11, e15)) = op2(e21, h5(e15))) | ~ (h5(op1(e10, e12)) = op2(e20, h5(e12))) | ~ (h5(op1(e11, e12)) = op2(e21, h5(e12))) | ~ (e20 = h5(op1(e12, e13))) | ~ (h5(e15) = h5(op1(e12, e14))) | ~ (h5(op1(e10, e13)) = op2(e20, h5(e13))) | ~ (h5(e15) = h5(op1(e11, e13))) | ~ (h5(op1(e12, e13)) = op2(h5(e13), h5(e12))) | ~ (e21 = h5(op1(e13, e14))) | ~ (op2(e20, e24) = h5(op1(e10, e14))) | ~ (h5(e12) = h5(op1(e11, e14))) | ~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f6254, f4941])).
fof(f6254, plain, (~ (h5(op1(e10, e12)) = op2(e20, h5(e12))) | ~ (h5(op1(e11, e12)) = op2(e21, h5(e12))) | ~ (e20 = h5(op1(e12, e13))) | ~ (h5(e15) = h5(op1(e12, e14))) | ~ (h5(op1(e10, e13)) = op2(e20, h5(e13))) | ~ (h5(e15) = h5(op1(e11, e13))) | ~ (h5(op1(e12, e13)) = op2(h5(e13), h5(e12))) | ~ (e21 = h5(op1(e13, e14))) | ~ (op2(e20, e24) = h5(op1(e10, e14))) | ~ (h5(e12) = h5(op1(e11, e14))) | ~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e11, e15)) = op2(h5(e11), h5(e15))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f6253, f1218])).
fof(f6253, plain, (~ (h5(op1(e12, e10)) = op2(e20, h5(e12))) | ~ (h5(op1(e11, e12)) = op2(e21, h5(e12))) | ~ (e20 = h5(op1(e12, e13))) | ~ (h5(e15) = h5(op1(e12, e14))) | ~ (h5(op1(e10, e13)) = op2(e20, h5(e13))) | ~ (h5(e15) = h5(op1(e11, e13))) | ~ (h5(op1(e12, e13)) = op2(h5(e13), h5(e12))) | ~ (e21 = h5(op1(e13, e14))) | ~ (op2(e20, e24) = h5(op1(e10, e14))) | ~ (h5(e12) = h5(op1(e11, e14))) | ~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e11, e15)) = op2(h5(e11), h5(e15))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f6252, f4909])).
fof(f4909, plain, (op2(e20, h5(e12)) = op2(h5(e12), e20)), inference(backward_demodulation, [], [f4662, f4888])).
fof(f4662, plain, (op2(e20, op2(e21, e24)) = op2(op2(e21, e24), e20)), inference(backward_demodulation, [], [f4438, f4661])).
fof(f4661, plain, (op2(e20, op2(e21, e24)) = op2(e21, op2(e20, e24))), inference(forward_demodulation, [], [f4639, f427])).
fof(f427, plain, (op2(op2(e20, e21), e24) = op2(e20, op2(e21, e24))), inference(cnf_transformation, [], [f6])).
fof(f4639, plain, (op2(op2(e20, e21), e24) = op2(e21, op2(e20, e24))), inference(backward_demodulation, [], [f457, f1248])).
fof(f457, plain, (op2(op2(e21, e20), e24) = op2(e21, op2(e20, e24))), inference(cnf_transformation, [], [f6])).
fof(f4438, plain, (op2(e21, op2(e20, e24)) = op2(op2(e21, e24), e20)), inference(backward_demodulation, [], [f477, f1266])).
fof(f1266, plain, (op2(e20, e24) = op2(e24, e20)), inference(cnf_transformation, [], [f23])).
fof(f477, plain, (op2(op2(e21, e24), e20) = op2(e21, op2(e24, e20))), inference(cnf_transformation, [], [f6])).
fof(f6252, plain, (~ (h5(op1(e12, e10)) = op2(h5(e12), e20)) | ~ (h5(op1(e11, e12)) = op2(e21, h5(e12))) | ~ (e20 = h5(op1(e12, e13))) | ~ (h5(e15) = h5(op1(e12, e14))) | ~ (h5(op1(e10, e13)) = op2(e20, h5(e13))) | ~ (h5(e15) = h5(op1(e11, e13))) | ~ (h5(op1(e12, e13)) = op2(h5(e13), h5(e12))) | ~ (e21 = h5(op1(e13, e14))) | ~ (op2(e20, e24) = h5(op1(e10, e14))) | ~ (h5(e12) = h5(op1(e11, e14))) | ~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e11, e15)) = op2(h5(e11), h5(e15))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f6251, f4945])).
fof(f6251, plain, (~ (h5(op1(e11, e12)) = op2(e21, h5(e12))) | ~ (e20 = h5(op1(e12, e13))) | ~ (h5(e15) = h5(op1(e12, e14))) | ~ (h5(op1(e10, e13)) = op2(e20, h5(e13))) | ~ (h5(e15) = h5(op1(e11, e13))) | ~ (h5(op1(e12, e13)) = op2(h5(e13), h5(e12))) | ~ (e21 = h5(op1(e13, e14))) | ~ (op2(e20, e24) = h5(op1(e10, e14))) | ~ (h5(e12) = h5(op1(e11, e14))) | ~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e15)) = op2(h5(e11), h5(e15))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f6250, f1219])).
fof(f1219, plain, (op1(e11, e12) = op1(e12, e11)), inference(cnf_transformation, [], [f22])).
fof(f6250, plain, (~ (h5(op1(e12, e11)) = op2(e21, h5(e12))) | ~ (e20 = h5(op1(e12, e13))) | ~ (h5(e15) = h5(op1(e12, e14))) | ~ (h5(op1(e10, e13)) = op2(e20, h5(e13))) | ~ (h5(e15) = h5(op1(e11, e13))) | ~ (h5(op1(e12, e13)) = op2(h5(e13), h5(e12))) | ~ (e21 = h5(op1(e13, e14))) | ~ (op2(e20, e24) = h5(op1(e10, e14))) | ~ (h5(e12) = h5(op1(e11, e14))) | ~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e15)) = op2(h5(e11), h5(e15))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f6249, f4939])).
fof(f4939, plain, (op2(e21, h5(e12)) = op2(h5(e12), e21)), inference(forward_demodulation, [], [f4911, f4910])).
fof(f4910, plain, (op2(e24, h2(e13)) = op2(e21, h5(e12))), inference(backward_demodulation, [], [f4752, f4888])).
fof(f4752, plain, (op2(e21, op2(e21, e24)) = op2(e24, h2(e13))), inference(backward_demodulation, [], [f4429, f1298])).
fof(f4429, plain, (op2(e21, op2(e21, e24)) = op2(e24, op2(e21, e21))), inference(forward_demodulation, [], [f4411, f4407])).
fof(f4407, plain, (op2(e21, op2(e21, e24)) = op2(op2(e21, e24), e21)), inference(backward_demodulation, [], [f478, f1267])).
fof(f478, plain, (op2(op2(e21, e24), e21) = op2(e21, op2(e24, e21))), inference(cnf_transformation, [], [f6])).
fof(f4411, plain, (op2(op2(e21, e24), e21) = op2(e24, op2(e21, e21))), inference(backward_demodulation, [], [f568, f1267])).
fof(f568, plain, (op2(op2(e24, e21), e21) = op2(e24, op2(e21, e21))), inference(cnf_transformation, [], [f6])).
fof(f4911, plain, (op2(e24, h2(e13)) = op2(h5(e12), e21)), inference(backward_demodulation, [], [f4765, f4888])).
fof(f4765, plain, (op2(op2(e21, e24), e21) = op2(e24, h2(e13))), inference(backward_demodulation, [], [f4760, f4764])).
fof(f4764, plain, (op2(h2(e13), e24) = op2(e24, h2(e13))), inference(backward_demodulation, [], [f4740, f4752])).
fof(f4740, plain, (op2(e21, op2(e21, e24)) = op2(h2(e13), e24)), inference(backward_demodulation, [], [f463, f1298])).
fof(f463, plain, (op2(op2(e21, e21), e24) = op2(e21, op2(e21, e24))), inference(cnf_transformation, [], [f6])).
fof(f4760, plain, (op2(op2(e21, e24), e21) = op2(h2(e13), e24)), inference(backward_demodulation, [], [f4407, f4740])).
fof(f6249, plain, (~ (h5(op1(e12, e11)) = op2(h5(e12), e21)) | ~ (e20 = h5(op1(e12, e13))) | ~ (h5(e15) = h5(op1(e12, e14))) | ~ (h5(op1(e10, e13)) = op2(e20, h5(e13))) | ~ (h5(e15) = h5(op1(e11, e13))) | ~ (h5(op1(e12, e13)) = op2(h5(e13), h5(e12))) | ~ (e21 = h5(op1(e13, e14))) | ~ (op2(e20, e24) = h5(op1(e10, e14))) | ~ (h5(e12) = h5(op1(e11, e14))) | ~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e15)) = op2(h5(e11), h5(e15))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f6248, f4941])).
fof(f6248, plain, (~ (e20 = h5(op1(e12, e13))) | ~ (h5(e15) = h5(op1(e12, e14))) | ~ (h5(op1(e10, e13)) = op2(e20, h5(e13))) | ~ (h5(e15) = h5(op1(e11, e13))) | ~ (h5(op1(e12, e13)) = op2(h5(e13), h5(e12))) | ~ (e21 = h5(op1(e13, e14))) | ~ (op2(e20, e24) = h5(op1(e10, e14))) | ~ (h5(e12) = h5(op1(e11, e14))) | ~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e15)) = op2(h5(e11), h5(e15))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f6247, f4915])).
fof(f6247, plain, (~ (h5(e15) = h5(op1(e12, e14))) | ~ (h5(op1(e10, e13)) = op2(e20, h5(e13))) | ~ (h5(e15) = h5(op1(e11, e13))) | ~ (h5(op1(e12, e13)) = op2(h5(e13), h5(e12))) | ~ (e21 = h5(op1(e13, e14))) | ~ (op2(e20, e24) = h5(op1(e10, e14))) | ~ (h5(e12) = h5(op1(e11, e14))) | ~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e15)) = op2(h5(e11), h5(e15))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(duplicate_literal_removal, [], [f6246])).
fof(f6246, plain, (~ (h5(e15) = h5(op1(e12, e14))) | ~ (h5(op1(e10, e13)) = op2(e20, h5(e13))) | ~ (h5(e15) = h5(op1(e11, e13))) | ~ (h5(op1(e12, e13)) = op2(h5(e13), h5(e12))) | ~ (e21 = h5(op1(e13, e14))) | ~ (op2(e20, e24) = h5(op1(e10, e14))) | ~ (h5(e12) = h5(op1(e11, e14))) | ~ (h5(e15) = h5(op1(e12, e14))) | ~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e15)) = op2(h5(e11), h5(e15))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f6245, f4912])).
fof(f6245, plain, (~ (h5(op1(e12, e14)) = op2(h5(e12), e24)) | ~ (h5(op1(e10, e13)) = op2(e20, h5(e13))) | ~ (h5(e15) = h5(op1(e11, e13))) | ~ (h5(op1(e12, e13)) = op2(h5(e13), h5(e12))) | ~ (e21 = h5(op1(e13, e14))) | ~ (op2(e20, e24) = h5(op1(e10, e14))) | ~ (h5(e12) = h5(op1(e11, e14))) | ~ (h5(e15) = h5(op1(e12, e14))) | ~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e15)) = op2(h5(e11), h5(e15))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f6244, f1312])).
fof(f6244, plain, (~ (h5(op1(e10, e13)) = op2(e20, h5(e13))) | ~ (h5(e15) = h5(op1(e11, e13))) | ~ (h5(op1(e12, e13)) = op2(h5(e13), h5(e12))) | ~ (e21 = h5(op1(e13, e14))) | ~ (op2(e20, e24) = h5(op1(e10, e14))) | ~ (h5(e12) = h5(op1(e11, e14))) | ~ (h5(e15) = h5(op1(e12, e14))) | ~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e15)) = op2(h5(e11), h5(e15))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f6243, f1224])).
fof(f6243, plain, (~ (h5(op1(e13, e10)) = op2(e20, h5(e13))) | ~ (h5(e15) = h5(op1(e11, e13))) | ~ (h5(op1(e12, e13)) = op2(h5(e13), h5(e12))) | ~ (e21 = h5(op1(e13, e14))) | ~ (op2(e20, e24) = h5(op1(e10, e14))) | ~ (h5(e12) = h5(op1(e11, e14))) | ~ (h5(e15) = h5(op1(e12, e14))) | ~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e15)) = op2(h5(e11), h5(e15))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f6242, f4879])).
fof(f4879, plain, (op2(e20, h5(e13)) = op2(h5(e13), e20)), inference(backward_demodulation, [], [f4468, f1316])).
fof(f4468, plain, (op2(e20, op2(e24, e24)) = op2(op2(e24, e24), e20)), inference(forward_demodulation, [], [f4447, f4465])).
fof(f4465, plain, (op2(e20, op2(e24, e24)) = op2(e24, op2(e20, e24))), inference(forward_demodulation, [], [f4445, f445])).
fof(f445, plain, (op2(op2(e20, e24), e24) = op2(e20, op2(e24, e24))), inference(cnf_transformation, [], [f6])).
fof(f4445, plain, (op2(op2(e20, e24), e24) = op2(e24, op2(e20, e24))), inference(backward_demodulation, [], [f565, f1266])).
fof(f565, plain, (op2(op2(e24, e20), e24) = op2(e24, op2(e20, e24))), inference(cnf_transformation, [], [f6])).
fof(f4447, plain, (op2(e24, op2(e20, e24)) = op2(op2(e24, e24), e20)), inference(backward_demodulation, [], [f585, f1266])).
fof(f585, plain, (op2(op2(e24, e24), e20) = op2(e24, op2(e24, e20))), inference(cnf_transformation, [], [f6])).
fof(f6242, plain, (~ (h5(op1(e13, e10)) = op2(h5(e13), e20)) | ~ (h5(e15) = h5(op1(e11, e13))) | ~ (h5(op1(e12, e13)) = op2(h5(e13), h5(e12))) | ~ (e21 = h5(op1(e13, e14))) | ~ (op2(e20, e24) = h5(op1(e10, e14))) | ~ (h5(e12) = h5(op1(e11, e14))) | ~ (h5(e15) = h5(op1(e12, e14))) | ~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e15)) = op2(h5(e11), h5(e15))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f6241, f4945])).
fof(f6241, plain, (~ (h5(e15) = h5(op1(e11, e13))) | ~ (h5(op1(e12, e13)) = op2(h5(e13), h5(e12))) | ~ (e21 = h5(op1(e13, e14))) | ~ (op2(e20, e24) = h5(op1(e10, e14))) | ~ (h5(e12) = h5(op1(e11, e14))) | ~ (h5(e15) = h5(op1(e12, e14))) | ~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e15)) = op2(h5(e11), h5(e15))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f6240, f1225])).
fof(f6240, plain, (~ (h5(e15) = h5(op1(e13, e11))) | ~ (h5(op1(e12, e13)) = op2(h5(e13), h5(e12))) | ~ (e21 = h5(op1(e13, e14))) | ~ (op2(e20, e24) = h5(op1(e10, e14))) | ~ (h5(e12) = h5(op1(e11, e14))) | ~ (h5(e15) = h5(op1(e12, e14))) | ~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e15)) = op2(h5(e11), h5(e15))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f6239, f4883])).
fof(f4883, plain, (h5(e15) = op2(h5(e13), e21)), inference(backward_demodulation, [], [f4860, f1316])).
fof(f4860, plain, (op2(op2(e24, e24), e21) = h5(e15)), inference(backward_demodulation, [], [f4435, f4859])).
fof(f4435, plain, (op2(e21, op2(e24, e24)) = op2(op2(e24, e24), e21)), inference(forward_demodulation, [], [f4416, f4432])).
fof(f4432, plain, (op2(e21, op2(e24, e24)) = op2(e24, op2(e21, e24))), inference(forward_demodulation, [], [f4414, f481])).
fof(f4414, plain, (op2(op2(e21, e24), e24) = op2(e24, op2(e21, e24))), inference(backward_demodulation, [], [f571, f1267])).
fof(f571, plain, (op2(op2(e24, e21), e24) = op2(e24, op2(e21, e24))), inference(cnf_transformation, [], [f6])).
fof(f4416, plain, (op2(e24, op2(e21, e24)) = op2(op2(e24, e24), e21)), inference(backward_demodulation, [], [f586, f1267])).
fof(f586, plain, (op2(op2(e24, e24), e21) = op2(e24, op2(e24, e21))), inference(cnf_transformation, [], [f6])).
fof(f6239, plain, (~ (h5(op1(e13, e11)) = op2(h5(e13), e21)) | ~ (h5(op1(e12, e13)) = op2(h5(e13), h5(e12))) | ~ (e21 = h5(op1(e13, e14))) | ~ (op2(e20, e24) = h5(op1(e10, e14))) | ~ (h5(e12) = h5(op1(e11, e14))) | ~ (h5(e15) = h5(op1(e12, e14))) | ~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e15)) = op2(h5(e11), h5(e15))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f6238, f4941])).
fof(f6238, plain, (~ (h5(op1(e12, e13)) = op2(h5(e13), h5(e12))) | ~ (e21 = h5(op1(e13, e14))) | ~ (op2(e20, e24) = h5(op1(e10, e14))) | ~ (h5(e12) = h5(op1(e11, e14))) | ~ (h5(e15) = h5(op1(e12, e14))) | ~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e15)) = op2(h5(e11), h5(e15))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f6237, f1226])).
fof(f6237, plain, (~ (e21 = h5(op1(e13, e14))) | ~ (op2(e20, e24) = h5(op1(e10, e14))) | ~ (h5(e12) = h5(op1(e11, e14))) | ~ (h5(e15) = h5(op1(e12, e14))) | ~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e15)) = op2(h5(e11), h5(e15))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(duplicate_literal_removal, [], [f6236])).
fof(f6236, plain, (~ (e21 = h5(op1(e13, e14))) | ~ (op2(e20, e24) = h5(op1(e10, e14))) | ~ (h5(e12) = h5(op1(e11, e14))) | ~ (h5(e15) = h5(op1(e12, e14))) | ~ (e21 = h5(op1(e13, e14))) | ~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e15)) = op2(h5(e11), h5(e15))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f6235, f4872])).
fof(f6235, plain, (~ (h5(op1(e13, e14)) = op2(h5(e13), e24)) | ~ (op2(e20, e24) = h5(op1(e10, e14))) | ~ (h5(e12) = h5(op1(e11, e14))) | ~ (h5(e15) = h5(op1(e12, e14))) | ~ (e21 = h5(op1(e13, e14))) | ~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e15)) = op2(h5(e11), h5(e15))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f6234, f1312])).
fof(f6234, plain, (~ (op2(e20, e24) = h5(op1(e10, e14))) | ~ (h5(e12) = h5(op1(e11, e14))) | ~ (h5(e15) = h5(op1(e12, e14))) | ~ (e21 = h5(op1(e13, e14))) | ~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e15)) = op2(h5(e11), h5(e15))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f6233, f1266])).
fof(f6233, plain, (~ (op2(e24, e20) = h5(op1(e10, e14))) | ~ (h5(e12) = h5(op1(e11, e14))) | ~ (h5(e15) = h5(op1(e12, e14))) | ~ (e21 = h5(op1(e13, e14))) | ~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e15)) = op2(h5(e11), h5(e15))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f6232, f1230])).
fof(f6232, plain, (~ (op2(e24, e20) = h5(op1(e14, e10))) | ~ (h5(e12) = h5(op1(e11, e14))) | ~ (h5(e15) = h5(op1(e12, e14))) | ~ (e21 = h5(op1(e13, e14))) | ~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e15)) = op2(h5(e11), h5(e15))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f6231, f1312])).
fof(f6231, plain, (~ (h5(op1(e14, e10)) = op2(h5(e14), e20)) | ~ (h5(e12) = h5(op1(e11, e14))) | ~ (h5(e15) = h5(op1(e12, e14))) | ~ (e21 = h5(op1(e13, e14))) | ~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e15)) = op2(h5(e11), h5(e15))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f6230, f4945])).
fof(f6230, plain, (~ (h5(e12) = h5(op1(e11, e14))) | ~ (h5(e15) = h5(op1(e12, e14))) | ~ (e21 = h5(op1(e13, e14))) | ~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e15)) = op2(h5(e11), h5(e15))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f6229, f4898])).
fof(f4898, plain, (op2(e24, e21) = h5(e12)), inference(backward_demodulation, [], [f1267, f4888])).
fof(f6229, plain, (~ (op2(e24, e21) = h5(op1(e11, e14))) | ~ (h5(e15) = h5(op1(e12, e14))) | ~ (e21 = h5(op1(e13, e14))) | ~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e15)) = op2(h5(e11), h5(e15))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f6228, f1231])).
fof(f6228, plain, (~ (op2(e24, e21) = h5(op1(e14, e11))) | ~ (h5(e15) = h5(op1(e12, e14))) | ~ (e21 = h5(op1(e13, e14))) | ~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e15)) = op2(h5(e11), h5(e15))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f6227, f1312])).
fof(f6227, plain, (~ (h5(op1(e14, e11)) = op2(h5(e14), e21)) | ~ (h5(e15) = h5(op1(e12, e14))) | ~ (e21 = h5(op1(e13, e14))) | ~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e15)) = op2(h5(e11), h5(e15))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f6226, f4941])).
fof(f6226, plain, (~ (h5(e15) = h5(op1(e12, e14))) | ~ (e21 = h5(op1(e13, e14))) | ~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e14, e11)) = op2(h5(e14), h5(e11))) | ~ (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e15)) = op2(h5(e11), h5(e15))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f6225, f1232])).
fof(f6225, plain, (~ (h5(e15) = h5(op1(e14, e12))) | ~ (e21 = h5(op1(e13, e14))) | ~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e14, e11)) = op2(h5(e14), h5(e11))) | ~ (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e15)) = op2(h5(e11), h5(e15))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f6224, f4913])).
fof(f4913, plain, (h5(e15) = op2(e24, h5(e12))), inference(backward_demodulation, [], [f4861, f4888])).
fof(f4861, plain, (op2(e24, op2(e21, e24)) = h5(e15)), inference(backward_demodulation, [], [f4432, f4859])).
fof(f6224, plain, (~ (h5(op1(e14, e12)) = op2(e24, h5(e12))) | ~ (e21 = h5(op1(e13, e14))) | ~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e14, e11)) = op2(h5(e14), h5(e11))) | ~ (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e15)) = op2(h5(e11), h5(e15))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f6223, f1312])).
fof(f6223, plain, (~ (e21 = h5(op1(e13, e14))) | ~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e14, e12)) = op2(h5(e14), h5(e12))) | ~ (h5(op1(e14, e11)) = op2(h5(e14), h5(e11))) | ~ (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e15)) = op2(h5(e11), h5(e15))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f6222, f1233])).
fof(f6222, plain, (~ (e21 = h5(op1(e14, e13))) | ~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e14, e12)) = op2(h5(e14), h5(e12))) | ~ (h5(op1(e14, e11)) = op2(h5(e14), h5(e11))) | ~ (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e15)) = op2(h5(e11), h5(e15))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f6221, f4880])).
fof(f6221, plain, (~ (h5(op1(e14, e13)) = op2(e24, h5(e13))) | ~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e14, e12)) = op2(h5(e14), h5(e12))) | ~ (h5(op1(e14, e11)) = op2(h5(e14), h5(e11))) | ~ (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e15)) = op2(h5(e11), h5(e15))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f6220, f1312])).
fof(f6220, plain, (~ (h5(e13) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e14, e13)) = op2(h5(e14), h5(e13))) | ~ (h5(op1(e14, e12)) = op2(h5(e14), h5(e12))) | ~ (h5(op1(e14, e11)) = op2(h5(e14), h5(e11))) | ~ (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e15)) = op2(h5(e11), h5(e15))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f6219, f1316])).
fof(f6219, plain, (~ (op2(e24, e24) = h5(op1(e14, e14))) | ~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e14, e13)) = op2(h5(e14), h5(e13))) | ~ (h5(op1(e14, e12)) = op2(h5(e14), h5(e12))) | ~ (h5(op1(e14, e11)) = op2(h5(e14), h5(e11))) | ~ (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e15)) = op2(h5(e11), h5(e15))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f6218, f1312])).
fof(f6218, plain, (~ (h5(op1(e14, e15)) = op2(e24, h5(e15))) | ~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e14, e14)) = op2(h5(e14), h5(e14))) | ~ (h5(op1(e14, e13)) = op2(h5(e14), h5(e13))) | ~ (h5(op1(e14, e12)) = op2(h5(e14), h5(e12))) | ~ (h5(op1(e14, e11)) = op2(h5(e14), h5(e11))) | ~ (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e15)) = op2(h5(e11), h5(e15))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f6217, f1312])).
fof(f6217, plain, (~ (h5(op1(e10, e15)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e14, e15)) = op2(h5(e14), h5(e15))) | ~ (h5(op1(e14, e14)) = op2(h5(e14), h5(e14))) | ~ (h5(op1(e14, e13)) = op2(h5(e14), h5(e13))) | ~ (h5(op1(e14, e12)) = op2(h5(e14), h5(e12))) | ~ (h5(op1(e14, e11)) = op2(h5(e14), h5(e11))) | ~ (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e15)) = op2(h5(e11), h5(e15))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f6216, f1236])).
fof(f1236, plain, (op1(e10, e15) = op1(e15, e10)), inference(cnf_transformation, [], [f22])).
fof(f6216, plain, (~ (h5(op1(e15, e10)) = op2(h5(e15), e20)) | ~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e14, e15)) = op2(h5(e14), h5(e15))) | ~ (h5(op1(e14, e14)) = op2(h5(e14), h5(e14))) | ~ (h5(op1(e14, e13)) = op2(h5(e14), h5(e13))) | ~ (h5(op1(e14, e12)) = op2(h5(e14), h5(e12))) | ~ (h5(op1(e14, e11)) = op2(h5(e14), h5(e11))) | ~ (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e15)) = op2(h5(e11), h5(e15))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f6215, f4945])).
fof(f6215, plain, (~ (h5(op1(e11, e15)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e15, e10)) = op2(h5(e15), h5(e10))) | ~ (h5(op1(e14, e15)) = op2(h5(e14), h5(e15))) | ~ (h5(op1(e14, e14)) = op2(h5(e14), h5(e14))) | ~ (h5(op1(e14, e13)) = op2(h5(e14), h5(e13))) | ~ (h5(op1(e14, e12)) = op2(h5(e14), h5(e12))) | ~ (h5(op1(e14, e11)) = op2(h5(e14), h5(e11))) | ~ (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e15)) = op2(h5(e11), h5(e15))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f6214, f1237])).
fof(f1237, plain, (op1(e11, e15) = op1(e15, e11)), inference(cnf_transformation, [], [f22])).
fof(f6214, plain, (~ (h5(op1(e15, e11)) = op2(h5(e15), e21)) | ~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e15, e10)) = op2(h5(e15), h5(e10))) | ~ (h5(op1(e14, e15)) = op2(h5(e14), h5(e15))) | ~ (h5(op1(e14, e14)) = op2(h5(e14), h5(e14))) | ~ (h5(op1(e14, e13)) = op2(h5(e14), h5(e13))) | ~ (h5(op1(e14, e12)) = op2(h5(e14), h5(e12))) | ~ (h5(op1(e14, e11)) = op2(h5(e14), h5(e11))) | ~ (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e15)) = op2(h5(e11), h5(e15))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f6213, f4941])).
fof(f6213, plain, (~ (h5(op1(e12, e15)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e15, e11)) = op2(h5(e15), h5(e11))) | ~ (h5(op1(e15, e10)) = op2(h5(e15), h5(e10))) | ~ (h5(op1(e14, e15)) = op2(h5(e14), h5(e15))) | ~ (h5(op1(e14, e14)) = op2(h5(e14), h5(e14))) | ~ (h5(op1(e14, e13)) = op2(h5(e14), h5(e13))) | ~ (h5(op1(e14, e12)) = op2(h5(e14), h5(e12))) | ~ (h5(op1(e14, e11)) = op2(h5(e14), h5(e11))) | ~ (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e15)) = op2(h5(e11), h5(e15))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f6212, f1238])).
fof(f6212, plain, (~ (h5(op1(e13, e15)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e15, e12)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e15, e11)) = op2(h5(e15), h5(e11))) | ~ (h5(op1(e15, e10)) = op2(h5(e15), h5(e10))) | ~ (h5(op1(e14, e15)) = op2(h5(e14), h5(e15))) | ~ (h5(op1(e14, e14)) = op2(h5(e14), h5(e14))) | ~ (h5(op1(e14, e13)) = op2(h5(e14), h5(e13))) | ~ (h5(op1(e14, e12)) = op2(h5(e14), h5(e12))) | ~ (h5(op1(e14, e11)) = op2(h5(e14), h5(e11))) | ~ (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e15)) = op2(h5(e11), h5(e15))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f6211, f1239])).
fof(f1239, plain, (op1(e13, e15) = op1(e15, e13)), inference(cnf_transformation, [], [f22])).
fof(f6211, plain, (~ (h5(op1(e14, e15)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e15, e13)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e15, e12)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e15, e11)) = op2(h5(e15), h5(e11))) | ~ (h5(op1(e15, e10)) = op2(h5(e15), h5(e10))) | ~ (h5(op1(e14, e15)) = op2(h5(e14), h5(e15))) | ~ (h5(op1(e14, e14)) = op2(h5(e14), h5(e14))) | ~ (h5(op1(e14, e13)) = op2(h5(e14), h5(e13))) | ~ (h5(op1(e14, e12)) = op2(h5(e14), h5(e12))) | ~ (h5(op1(e14, e11)) = op2(h5(e14), h5(e11))) | ~ (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e15)) = op2(h5(e11), h5(e15))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f6210, f1240])).
fof(f1240, plain, (op1(e14, e15) = op1(e15, e14)), inference(cnf_transformation, [], [f22])).
fof(f6210, plain, (~ (h5(op1(e15, e14)) = op2(h5(e15), e24)) | ~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e15, e13)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e15, e12)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e15, e11)) = op2(h5(e15), h5(e11))) | ~ (h5(op1(e15, e10)) = op2(h5(e15), h5(e10))) | ~ (h5(op1(e14, e15)) = op2(h5(e14), h5(e15))) | ~ (h5(op1(e14, e14)) = op2(h5(e14), h5(e14))) | ~ (h5(op1(e14, e13)) = op2(h5(e14), h5(e13))) | ~ (h5(op1(e14, e12)) = op2(h5(e14), h5(e12))) | ~ (h5(op1(e14, e11)) = op2(h5(e14), h5(e11))) | ~ (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e15)) = op2(h5(e11), h5(e15))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f1533, f1312])).
fof(f1533, plain, (~ (e25 = h5(e15)) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e15, e14)) = op2(h5(e15), h5(e14))) | ~ (h5(op1(e15, e13)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e15, e12)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e15, e11)) = op2(h5(e15), h5(e11))) | ~ (h5(op1(e15, e10)) = op2(h5(e15), h5(e10))) | ~ (h5(op1(e14, e15)) = op2(h5(e14), h5(e15))) | ~ (h5(op1(e14, e14)) = op2(h5(e14), h5(e14))) | ~ (h5(op1(e14, e13)) = op2(h5(e14), h5(e13))) | ~ (h5(op1(e14, e12)) = op2(h5(e14), h5(e12))) | ~ (h5(op1(e14, e11)) = op2(h5(e14), h5(e11))) | ~ (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e15)) = op2(h5(e11), h5(e15))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(cnf_transformation, [], [f67])).
fof(f67, plain, (((~ (e25 = h6(e15)) & ~ (e25 = h6(e14)) & ~ (e25 = h6(e13)) & ~ (e25 = h6(e12)) & ~ (e25 = h6(e11)) & ~ (e25 = h6(e10))) | sP29 | sP28 | sP27 | sP26 | sP25 | ~ (h6(op1(e15, e15)) = op2(h6(e15), h6(e15))) | ~ (h6(op1(e15, e14)) = op2(h6(e15), h6(e14))) | ~ (h6(op1(e15, e13)) = op2(h6(e15), h6(e13))) | ~ (h6(op1(e15, e12)) = op2(h6(e15), h6(e12))) | ~ (h6(op1(e15, e11)) = op2(h6(e15), h6(e11))) | ~ (h6(op1(e15, e10)) = op2(h6(e15), h6(e10))) | ~ (h6(op1(e14, e15)) = op2(h6(e14), h6(e15))) | ~ (h6(op1(e14, e14)) = op2(h6(e14), h6(e14))) | ~ (h6(op1(e14, e13)) = op2(h6(e14), h6(e13))) | ~ (h6(op1(e14, e12)) = op2(h6(e14), h6(e12))) | ~ (h6(op1(e14, e11)) = op2(h6(e14), h6(e11))) | ~ (h6(op1(e14, e10)) = op2(h6(e14), h6(e10))) | ~ (h6(op1(e13, e15)) = op2(h6(e13), h6(e15))) | ~ (h6(op1(e13, e14)) = op2(h6(e13), h6(e14))) | ~ (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) | ~ (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e15)) = op2(h6(e12), h6(e15))) | ~ (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))) & ((~ (e25 = h5(e15)) & ~ (e25 = h5(e14)) & ~ (e25 = h5(e13)) & ~ (e25 = h5(e12)) & ~ (e25 = h5(e11)) & ~ (e25 = h5(e10))) | sP24 | sP23 | sP22 | sP21 | sP20 | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e15, e14)) = op2(h5(e15), h5(e14))) | ~ (h5(op1(e15, e13)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e15, e12)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e15, e11)) = op2(h5(e15), h5(e11))) | ~ (h5(op1(e15, e10)) = op2(h5(e15), h5(e10))) | ~ (h5(op1(e14, e15)) = op2(h5(e14), h5(e15))) | ~ (h5(op1(e14, e14)) = op2(h5(e14), h5(e14))) | ~ (h5(op1(e14, e13)) = op2(h5(e14), h5(e13))) | ~ (h5(op1(e14, e12)) = op2(h5(e14), h5(e12))) | ~ (h5(op1(e14, e11)) = op2(h5(e14), h5(e11))) | ~ (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e15)) = op2(h5(e11), h5(e15))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))) & ((~ (e25 = h4(e15)) & ~ (e25 = h4(e14)) & ~ (e25 = h4(e13)) & ~ (e25 = h4(e12)) & ~ (e25 = h4(e11)) & ~ (e25 = h4(e10))) | sP19 | sP18 | sP17 | sP16 | sP15 | ~ (h4(op1(e15, e15)) = op2(h4(e15), h4(e15))) | ~ (h4(op1(e15, e14)) = op2(h4(e15), h4(e14))) | ~ (h4(op1(e15, e13)) = op2(h4(e15), h4(e13))) | ~ (h4(op1(e15, e12)) = op2(h4(e15), h4(e12))) | ~ (h4(op1(e15, e11)) = op2(h4(e15), h4(e11))) | ~ (h4(op1(e15, e10)) = op2(h4(e15), h4(e10))) | ~ (h4(op1(e14, e15)) = op2(h4(e14), h4(e15))) | ~ (h4(op1(e14, e14)) = op2(h4(e14), h4(e14))) | ~ (h4(op1(e14, e13)) = op2(h4(e14), h4(e13))) | ~ (h4(op1(e14, e12)) = op2(h4(e14), h4(e12))) | ~ (h4(op1(e14, e11)) = op2(h4(e14), h4(e11))) | ~ (h4(op1(e14, e10)) = op2(h4(e14), h4(e10))) | ~ (h4(op1(e13, e15)) = op2(h4(e13), h4(e15))) | ~ (h4(op1(e13, e14)) = op2(h4(e13), h4(e14))) | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e15)) = op2(h4(e12), h4(e15))) | ~ (h4(op1(e12, e14)) = op2(h4(e12), h4(e14))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e15)) = op2(h4(e11), h4(e15))) | ~ (h4(op1(e11, e14)) = op2(h4(e11), h4(e14))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e15)) = op2(h4(e10), h4(e15))) | ~ (h4(op1(e10, e14)) = op2(h4(e10), h4(e14))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) & ((~ (e25 = h3(e15)) & ~ (e25 = h3(e14)) & ~ (e25 = h3(e13)) & ~ (e25 = h3(e12)) & ~ (e25 = h3(e11)) & ~ (e25 = h3(e10))) | sP14 | sP13 | sP12 | sP11 | sP10 | ~ (h3(op1(e15, e15)) = op2(h3(e15), h3(e15))) | ~ (h3(op1(e15, e14)) = op2(h3(e15), h3(e14))) | ~ (h3(op1(e15, e13)) = op2(h3(e15), h3(e13))) | ~ (h3(op1(e15, e12)) = op2(h3(e15), h3(e12))) | ~ (h3(op1(e15, e11)) = op2(h3(e15), h3(e11))) | ~ (h3(op1(e15, e10)) = op2(h3(e15), h3(e10))) | ~ (h3(op1(e14, e15)) = op2(h3(e14), h3(e15))) | ~ (h3(op1(e14, e14)) = op2(h3(e14), h3(e14))) | ~ (h3(op1(e14, e13)) = op2(h3(e14), h3(e13))) | ~ (h3(op1(e14, e12)) = op2(h3(e14), h3(e12))) | ~ (h3(op1(e14, e11)) = op2(h3(e14), h3(e11))) | ~ (h3(op1(e14, e10)) = op2(h3(e14), h3(e10))) | ~ (h3(op1(e13, e15)) = op2(h3(e13), h3(e15))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), h3(e14))) | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e15)) = op2(h3(e12), h3(e15))) | ~ (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e15)) = op2(h3(e11), h3(e15))) | ~ (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e15)) = op2(h3(e10), h3(e15))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) & ((~ (e25 = h2(e15)) & ~ (e25 = h2(e14)) & ~ (e25 = h2(e13)) & ~ (e25 = h2(e12)) & ~ (e25 = h2(e11)) & ~ (e25 = h2(e10))) | sP9 | sP8 | sP7 | sP6 | sP5 | ~ (h2(op1(e15, e15)) = op2(h2(e15), h2(e15))) | ~ (h2(op1(e15, e14)) = op2(h2(e15), h2(e14))) | ~ (h2(op1(e15, e13)) = op2(h2(e15), h2(e13))) | ~ (h2(op1(e15, e12)) = op2(h2(e15), h2(e12))) | ~ (h2(op1(e15, e11)) = op2(h2(e15), h2(e11))) | ~ (h2(op1(e15, e10)) = op2(h2(e15), h2(e10))) | ~ (h2(op1(e14, e15)) = op2(h2(e14), h2(e15))) | ~ (h2(op1(e14, e14)) = op2(h2(e14), h2(e14))) | ~ (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) | ~ (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) | ~ (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) | ~ (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) | ~ (h2(op1(e13, e15)) = op2(h2(e13), h2(e15))) | ~ (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e15)) = op2(h2(e12), h2(e15))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e15)) = op2(h2(e11), h2(e15))) | ~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e15)) = op2(h2(e10), h2(e15))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) & ((~ (e25 = h1(e15)) & ~ (e25 = h1(e14)) & ~ (e25 = h1(e13)) & ~ (e25 = h1(e12)) & ~ (e25 = h1(e11)) & ~ (e25 = h1(e10))) | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (h1(op1(e15, e15)) = op2(h1(e15), h1(e15))) | ~ (h1(op1(e15, e14)) = op2(h1(e15), h1(e14))) | ~ (h1(op1(e15, e13)) = op2(h1(e15), h1(e13))) | ~ (h1(op1(e15, e12)) = op2(h1(e15), h1(e12))) | ~ (h1(op1(e15, e11)) = op2(h1(e15), h1(e11))) | ~ (h1(op1(e15, e10)) = op2(h1(e15), h1(e10))) | ~ (h1(op1(e14, e15)) = op2(h1(e14), h1(e15))) | ~ (h1(op1(e14, e14)) = op2(h1(e14), h1(e14))) | ~ (h1(op1(e14, e13)) = op2(h1(e14), h1(e13))) | ~ (h1(op1(e14, e12)) = op2(h1(e14), h1(e12))) | ~ (h1(op1(e14, e11)) = op2(h1(e14), h1(e11))) | ~ (h1(op1(e14, e10)) = op2(h1(e14), h1(e10))) | ~ (h1(op1(e13, e15)) = op2(h1(e13), h1(e15))) | ~ (h1(op1(e13, e14)) = op2(h1(e13), h1(e14))) | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e15)) = op2(h1(e12), h1(e15))) | ~ (h1(op1(e12, e14)) = op2(h1(e12), h1(e14))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e15)) = op2(h1(e11), h1(e15))) | ~ (h1(op1(e11, e14)) = op2(h1(e11), h1(e14))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e15)) = op2(h1(e10), h1(e15))) | ~ (h1(op1(e10, e14)) = op2(h1(e10), h1(e14))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(definition_folding, [], [f36, e66, e65, e64, e63, e62, e61, e60, e59, e58, e57, e56, e55, e54, e53, e52, e51, e50, e49, e48, e47, e46, e45, e44, e43, e42, e41, e40, e39, e38, e37])).
fof(f37, plain, ((~ (e20 = h1(e15)) & ~ (e20 = h1(e14)) & ~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10))) | ~ sP0), inference(usedef, [], [e37])).
fof(e37, plain, (sP0 <=> (~ (e20 = h1(e15)) & ~ (e20 = h1(e14)) & ~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f38, plain, ((~ (e21 = h1(e15)) & ~ (e21 = h1(e14)) & ~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10))) | ~ sP1), inference(usedef, [], [e38])).
fof(e38, plain, (sP1 <=> (~ (e21 = h1(e15)) & ~ (e21 = h1(e14)) & ~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f39, plain, ((~ (e22 = h1(e15)) & ~ (e22 = h1(e14)) & ~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10))) | ~ sP2), inference(usedef, [], [e39])).
fof(e39, plain, (sP2 <=> (~ (e22 = h1(e15)) & ~ (e22 = h1(e14)) & ~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f40, plain, ((~ (e23 = h1(e15)) & ~ (e23 = h1(e14)) & ~ (e23 = h1(e13)) & ~ (e23 = h1(e12)) & ~ (e23 = h1(e11)) & ~ (e23 = h1(e10))) | ~ sP3), inference(usedef, [], [e40])).
fof(e40, plain, (sP3 <=> (~ (e23 = h1(e15)) & ~ (e23 = h1(e14)) & ~ (e23 = h1(e13)) & ~ (e23 = h1(e12)) & ~ (e23 = h1(e11)) & ~ (e23 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f41, plain, ((~ (e24 = h1(e15)) & ~ (e24 = h1(e14)) & ~ (e24 = h1(e13)) & ~ (e24 = h1(e12)) & ~ (e24 = h1(e11)) & ~ (e24 = h1(e10))) | ~ sP4), inference(usedef, [], [e41])).
fof(e41, plain, (sP4 <=> (~ (e24 = h1(e15)) & ~ (e24 = h1(e14)) & ~ (e24 = h1(e13)) & ~ (e24 = h1(e12)) & ~ (e24 = h1(e11)) & ~ (e24 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f42, plain, ((~ (e20 = h2(e15)) & ~ (e20 = h2(e14)) & ~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10))) | ~ sP5), inference(usedef, [], [e42])).
fof(e42, plain, (sP5 <=> (~ (e20 = h2(e15)) & ~ (e20 = h2(e14)) & ~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f43, plain, ((~ (e21 = h2(e15)) & ~ (e21 = h2(e14)) & ~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10))) | ~ sP6), inference(usedef, [], [e43])).
fof(e43, plain, (sP6 <=> (~ (e21 = h2(e15)) & ~ (e21 = h2(e14)) & ~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f44, plain, ((~ (e22 = h2(e15)) & ~ (e22 = h2(e14)) & ~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10))) | ~ sP7), inference(usedef, [], [e44])).
fof(e44, plain, (sP7 <=> (~ (e22 = h2(e15)) & ~ (e22 = h2(e14)) & ~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f45, plain, ((~ (e23 = h2(e15)) & ~ (e23 = h2(e14)) & ~ (e23 = h2(e13)) & ~ (e23 = h2(e12)) & ~ (e23 = h2(e11)) & ~ (e23 = h2(e10))) | ~ sP8), inference(usedef, [], [e45])).
fof(e45, plain, (sP8 <=> (~ (e23 = h2(e15)) & ~ (e23 = h2(e14)) & ~ (e23 = h2(e13)) & ~ (e23 = h2(e12)) & ~ (e23 = h2(e11)) & ~ (e23 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f46, plain, ((~ (e24 = h2(e15)) & ~ (e24 = h2(e14)) & ~ (e24 = h2(e13)) & ~ (e24 = h2(e12)) & ~ (e24 = h2(e11)) & ~ (e24 = h2(e10))) | ~ sP9), inference(usedef, [], [e46])).
fof(e46, plain, (sP9 <=> (~ (e24 = h2(e15)) & ~ (e24 = h2(e14)) & ~ (e24 = h2(e13)) & ~ (e24 = h2(e12)) & ~ (e24 = h2(e11)) & ~ (e24 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f47, plain, ((~ (e20 = h3(e15)) & ~ (e20 = h3(e14)) & ~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10))) | ~ sP10), inference(usedef, [], [e47])).
fof(e47, plain, (sP10 <=> (~ (e20 = h3(e15)) & ~ (e20 = h3(e14)) & ~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f48, plain, ((~ (e21 = h3(e15)) & ~ (e21 = h3(e14)) & ~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10))) | ~ sP11), inference(usedef, [], [e48])).
fof(e48, plain, (sP11 <=> (~ (e21 = h3(e15)) & ~ (e21 = h3(e14)) & ~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f49, plain, ((~ (e22 = h3(e15)) & ~ (e22 = h3(e14)) & ~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10))) | ~ sP12), inference(usedef, [], [e49])).
fof(e49, plain, (sP12 <=> (~ (e22 = h3(e15)) & ~ (e22 = h3(e14)) & ~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f50, plain, ((~ (e23 = h3(e15)) & ~ (e23 = h3(e14)) & ~ (e23 = h3(e13)) & ~ (e23 = h3(e12)) & ~ (e23 = h3(e11)) & ~ (e23 = h3(e10))) | ~ sP13), inference(usedef, [], [e50])).
fof(e50, plain, (sP13 <=> (~ (e23 = h3(e15)) & ~ (e23 = h3(e14)) & ~ (e23 = h3(e13)) & ~ (e23 = h3(e12)) & ~ (e23 = h3(e11)) & ~ (e23 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f51, plain, ((~ (e24 = h3(e15)) & ~ (e24 = h3(e14)) & ~ (e24 = h3(e13)) & ~ (e24 = h3(e12)) & ~ (e24 = h3(e11)) & ~ (e24 = h3(e10))) | ~ sP14), inference(usedef, [], [e51])).
fof(e51, plain, (sP14 <=> (~ (e24 = h3(e15)) & ~ (e24 = h3(e14)) & ~ (e24 = h3(e13)) & ~ (e24 = h3(e12)) & ~ (e24 = h3(e11)) & ~ (e24 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f52, plain, ((~ (e20 = h4(e15)) & ~ (e20 = h4(e14)) & ~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10))) | ~ sP15), inference(usedef, [], [e52])).
fof(e52, plain, (sP15 <=> (~ (e20 = h4(e15)) & ~ (e20 = h4(e14)) & ~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f53, plain, ((~ (e21 = h4(e15)) & ~ (e21 = h4(e14)) & ~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10))) | ~ sP16), inference(usedef, [], [e53])).
fof(e53, plain, (sP16 <=> (~ (e21 = h4(e15)) & ~ (e21 = h4(e14)) & ~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f54, plain, ((~ (e22 = h4(e15)) & ~ (e22 = h4(e14)) & ~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10))) | ~ sP17), inference(usedef, [], [e54])).
fof(e54, plain, (sP17 <=> (~ (e22 = h4(e15)) & ~ (e22 = h4(e14)) & ~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f55, plain, ((~ (e23 = h4(e15)) & ~ (e23 = h4(e14)) & ~ (e23 = h4(e13)) & ~ (e23 = h4(e12)) & ~ (e23 = h4(e11)) & ~ (e23 = h4(e10))) | ~ sP18), inference(usedef, [], [e55])).
fof(e55, plain, (sP18 <=> (~ (e23 = h4(e15)) & ~ (e23 = h4(e14)) & ~ (e23 = h4(e13)) & ~ (e23 = h4(e12)) & ~ (e23 = h4(e11)) & ~ (e23 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f56, plain, ((~ (e24 = h4(e15)) & ~ (e24 = h4(e14)) & ~ (e24 = h4(e13)) & ~ (e24 = h4(e12)) & ~ (e24 = h4(e11)) & ~ (e24 = h4(e10))) | ~ sP19), inference(usedef, [], [e56])).
fof(e56, plain, (sP19 <=> (~ (e24 = h4(e15)) & ~ (e24 = h4(e14)) & ~ (e24 = h4(e13)) & ~ (e24 = h4(e12)) & ~ (e24 = h4(e11)) & ~ (e24 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP19])])).
fof(f57, plain, ((~ (e20 = h5(e15)) & ~ (e20 = h5(e14)) & ~ (e20 = h5(e13)) & ~ (e20 = h5(e12)) & ~ (e20 = h5(e11)) & ~ (e20 = h5(e10))) | ~ sP20), inference(usedef, [], [e57])).
fof(e57, plain, (sP20 <=> (~ (e20 = h5(e15)) & ~ (e20 = h5(e14)) & ~ (e20 = h5(e13)) & ~ (e20 = h5(e12)) & ~ (e20 = h5(e11)) & ~ (e20 = h5(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP20])])).
fof(f58, plain, ((~ (e21 = h5(e15)) & ~ (e21 = h5(e14)) & ~ (e21 = h5(e13)) & ~ (e21 = h5(e12)) & ~ (e21 = h5(e11)) & ~ (e21 = h5(e10))) | ~ sP21), inference(usedef, [], [e58])).
fof(e58, plain, (sP21 <=> (~ (e21 = h5(e15)) & ~ (e21 = h5(e14)) & ~ (e21 = h5(e13)) & ~ (e21 = h5(e12)) & ~ (e21 = h5(e11)) & ~ (e21 = h5(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP21])])).
fof(f59, plain, ((~ (e22 = h5(e15)) & ~ (e22 = h5(e14)) & ~ (e22 = h5(e13)) & ~ (e22 = h5(e12)) & ~ (e22 = h5(e11)) & ~ (e22 = h5(e10))) | ~ sP22), inference(usedef, [], [e59])).
fof(e59, plain, (sP22 <=> (~ (e22 = h5(e15)) & ~ (e22 = h5(e14)) & ~ (e22 = h5(e13)) & ~ (e22 = h5(e12)) & ~ (e22 = h5(e11)) & ~ (e22 = h5(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP22])])).
fof(f60, plain, ((~ (e23 = h5(e15)) & ~ (e23 = h5(e14)) & ~ (e23 = h5(e13)) & ~ (e23 = h5(e12)) & ~ (e23 = h5(e11)) & ~ (e23 = h5(e10))) | ~ sP23), inference(usedef, [], [e60])).
fof(e60, plain, (sP23 <=> (~ (e23 = h5(e15)) & ~ (e23 = h5(e14)) & ~ (e23 = h5(e13)) & ~ (e23 = h5(e12)) & ~ (e23 = h5(e11)) & ~ (e23 = h5(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP23])])).
fof(f61, plain, ((~ (e24 = h5(e15)) & ~ (e24 = h5(e14)) & ~ (e24 = h5(e13)) & ~ (e24 = h5(e12)) & ~ (e24 = h5(e11)) & ~ (e24 = h5(e10))) | ~ sP24), inference(usedef, [], [e61])).
fof(e61, plain, (sP24 <=> (~ (e24 = h5(e15)) & ~ (e24 = h5(e14)) & ~ (e24 = h5(e13)) & ~ (e24 = h5(e12)) & ~ (e24 = h5(e11)) & ~ (e24 = h5(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP24])])).
fof(f62, plain, ((~ (e20 = h6(e15)) & ~ (e20 = h6(e14)) & ~ (e20 = h6(e13)) & ~ (e20 = h6(e12)) & ~ (e20 = h6(e11)) & ~ (e20 = h6(e10))) | ~ sP25), inference(usedef, [], [e62])).
fof(e62, plain, (sP25 <=> (~ (e20 = h6(e15)) & ~ (e20 = h6(e14)) & ~ (e20 = h6(e13)) & ~ (e20 = h6(e12)) & ~ (e20 = h6(e11)) & ~ (e20 = h6(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP25])])).
fof(f63, plain, ((~ (e21 = h6(e15)) & ~ (e21 = h6(e14)) & ~ (e21 = h6(e13)) & ~ (e21 = h6(e12)) & ~ (e21 = h6(e11)) & ~ (e21 = h6(e10))) | ~ sP26), inference(usedef, [], [e63])).
fof(e63, plain, (sP26 <=> (~ (e21 = h6(e15)) & ~ (e21 = h6(e14)) & ~ (e21 = h6(e13)) & ~ (e21 = h6(e12)) & ~ (e21 = h6(e11)) & ~ (e21 = h6(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP26])])).
fof(f64, plain, ((~ (e22 = h6(e15)) & ~ (e22 = h6(e14)) & ~ (e22 = h6(e13)) & ~ (e22 = h6(e12)) & ~ (e22 = h6(e11)) & ~ (e22 = h6(e10))) | ~ sP27), inference(usedef, [], [e64])).
fof(e64, plain, (sP27 <=> (~ (e22 = h6(e15)) & ~ (e22 = h6(e14)) & ~ (e22 = h6(e13)) & ~ (e22 = h6(e12)) & ~ (e22 = h6(e11)) & ~ (e22 = h6(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP27])])).
fof(f65, plain, ((~ (e23 = h6(e15)) & ~ (e23 = h6(e14)) & ~ (e23 = h6(e13)) & ~ (e23 = h6(e12)) & ~ (e23 = h6(e11)) & ~ (e23 = h6(e10))) | ~ sP28), inference(usedef, [], [e65])).
fof(e65, plain, (sP28 <=> (~ (e23 = h6(e15)) & ~ (e23 = h6(e14)) & ~ (e23 = h6(e13)) & ~ (e23 = h6(e12)) & ~ (e23 = h6(e11)) & ~ (e23 = h6(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP28])])).
fof(f66, plain, ((~ (e24 = h6(e15)) & ~ (e24 = h6(e14)) & ~ (e24 = h6(e13)) & ~ (e24 = h6(e12)) & ~ (e24 = h6(e11)) & ~ (e24 = h6(e10))) | ~ sP29), inference(usedef, [], [e66])).
fof(e66, plain, (sP29 <=> (~ (e24 = h6(e15)) & ~ (e24 = h6(e14)) & ~ (e24 = h6(e13)) & ~ (e24 = h6(e12)) & ~ (e24 = h6(e11)) & ~ (e24 = h6(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP29])])).
fof(f36, plain, (((~ (e25 = h6(e15)) & ~ (e25 = h6(e14)) & ~ (e25 = h6(e13)) & ~ (e25 = h6(e12)) & ~ (e25 = h6(e11)) & ~ (e25 = h6(e10))) | (~ (e24 = h6(e15)) & ~ (e24 = h6(e14)) & ~ (e24 = h6(e13)) & ~ (e24 = h6(e12)) & ~ (e24 = h6(e11)) & ~ (e24 = h6(e10))) | (~ (e23 = h6(e15)) & ~ (e23 = h6(e14)) & ~ (e23 = h6(e13)) & ~ (e23 = h6(e12)) & ~ (e23 = h6(e11)) & ~ (e23 = h6(e10))) | (~ (e22 = h6(e15)) & ~ (e22 = h6(e14)) & ~ (e22 = h6(e13)) & ~ (e22 = h6(e12)) & ~ (e22 = h6(e11)) & ~ (e22 = h6(e10))) | (~ (e21 = h6(e15)) & ~ (e21 = h6(e14)) & ~ (e21 = h6(e13)) & ~ (e21 = h6(e12)) & ~ (e21 = h6(e11)) & ~ (e21 = h6(e10))) | (~ (e20 = h6(e15)) & ~ (e20 = h6(e14)) & ~ (e20 = h6(e13)) & ~ (e20 = h6(e12)) & ~ (e20 = h6(e11)) & ~ (e20 = h6(e10))) | ~ (h6(op1(e15, e15)) = op2(h6(e15), h6(e15))) | ~ (h6(op1(e15, e14)) = op2(h6(e15), h6(e14))) | ~ (h6(op1(e15, e13)) = op2(h6(e15), h6(e13))) | ~ (h6(op1(e15, e12)) = op2(h6(e15), h6(e12))) | ~ (h6(op1(e15, e11)) = op2(h6(e15), h6(e11))) | ~ (h6(op1(e15, e10)) = op2(h6(e15), h6(e10))) | ~ (h6(op1(e14, e15)) = op2(h6(e14), h6(e15))) | ~ (h6(op1(e14, e14)) = op2(h6(e14), h6(e14))) | ~ (h6(op1(e14, e13)) = op2(h6(e14), h6(e13))) | ~ (h6(op1(e14, e12)) = op2(h6(e14), h6(e12))) | ~ (h6(op1(e14, e11)) = op2(h6(e14), h6(e11))) | ~ (h6(op1(e14, e10)) = op2(h6(e14), h6(e10))) | ~ (h6(op1(e13, e15)) = op2(h6(e13), h6(e15))) | ~ (h6(op1(e13, e14)) = op2(h6(e13), h6(e14))) | ~ (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) | ~ (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e15)) = op2(h6(e12), h6(e15))) | ~ (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))) & ((~ (e25 = h5(e15)) & ~ (e25 = h5(e14)) & ~ (e25 = h5(e13)) & ~ (e25 = h5(e12)) & ~ (e25 = h5(e11)) & ~ (e25 = h5(e10))) | (~ (e24 = h5(e15)) & ~ (e24 = h5(e14)) & ~ (e24 = h5(e13)) & ~ (e24 = h5(e12)) & ~ (e24 = h5(e11)) & ~ (e24 = h5(e10))) | (~ (e23 = h5(e15)) & ~ (e23 = h5(e14)) & ~ (e23 = h5(e13)) & ~ (e23 = h5(e12)) & ~ (e23 = h5(e11)) & ~ (e23 = h5(e10))) | (~ (e22 = h5(e15)) & ~ (e22 = h5(e14)) & ~ (e22 = h5(e13)) & ~ (e22 = h5(e12)) & ~ (e22 = h5(e11)) & ~ (e22 = h5(e10))) | (~ (e21 = h5(e15)) & ~ (e21 = h5(e14)) & ~ (e21 = h5(e13)) & ~ (e21 = h5(e12)) & ~ (e21 = h5(e11)) & ~ (e21 = h5(e10))) | (~ (e20 = h5(e15)) & ~ (e20 = h5(e14)) & ~ (e20 = h5(e13)) & ~ (e20 = h5(e12)) & ~ (e20 = h5(e11)) & ~ (e20 = h5(e10))) | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e15, e14)) = op2(h5(e15), h5(e14))) | ~ (h5(op1(e15, e13)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e15, e12)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e15, e11)) = op2(h5(e15), h5(e11))) | ~ (h5(op1(e15, e10)) = op2(h5(e15), h5(e10))) | ~ (h5(op1(e14, e15)) = op2(h5(e14), h5(e15))) | ~ (h5(op1(e14, e14)) = op2(h5(e14), h5(e14))) | ~ (h5(op1(e14, e13)) = op2(h5(e14), h5(e13))) | ~ (h5(op1(e14, e12)) = op2(h5(e14), h5(e12))) | ~ (h5(op1(e14, e11)) = op2(h5(e14), h5(e11))) | ~ (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e15)) = op2(h5(e11), h5(e15))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))) & ((~ (e25 = h4(e15)) & ~ (e25 = h4(e14)) & ~ (e25 = h4(e13)) & ~ (e25 = h4(e12)) & ~ (e25 = h4(e11)) & ~ (e25 = h4(e10))) | (~ (e24 = h4(e15)) & ~ (e24 = h4(e14)) & ~ (e24 = h4(e13)) & ~ (e24 = h4(e12)) & ~ (e24 = h4(e11)) & ~ (e24 = h4(e10))) | (~ (e23 = h4(e15)) & ~ (e23 = h4(e14)) & ~ (e23 = h4(e13)) & ~ (e23 = h4(e12)) & ~ (e23 = h4(e11)) & ~ (e23 = h4(e10))) | (~ (e22 = h4(e15)) & ~ (e22 = h4(e14)) & ~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10))) | (~ (e21 = h4(e15)) & ~ (e21 = h4(e14)) & ~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10))) | (~ (e20 = h4(e15)) & ~ (e20 = h4(e14)) & ~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10))) | ~ (h4(op1(e15, e15)) = op2(h4(e15), h4(e15))) | ~ (h4(op1(e15, e14)) = op2(h4(e15), h4(e14))) | ~ (h4(op1(e15, e13)) = op2(h4(e15), h4(e13))) | ~ (h4(op1(e15, e12)) = op2(h4(e15), h4(e12))) | ~ (h4(op1(e15, e11)) = op2(h4(e15), h4(e11))) | ~ (h4(op1(e15, e10)) = op2(h4(e15), h4(e10))) | ~ (h4(op1(e14, e15)) = op2(h4(e14), h4(e15))) | ~ (h4(op1(e14, e14)) = op2(h4(e14), h4(e14))) | ~ (h4(op1(e14, e13)) = op2(h4(e14), h4(e13))) | ~ (h4(op1(e14, e12)) = op2(h4(e14), h4(e12))) | ~ (h4(op1(e14, e11)) = op2(h4(e14), h4(e11))) | ~ (h4(op1(e14, e10)) = op2(h4(e14), h4(e10))) | ~ (h4(op1(e13, e15)) = op2(h4(e13), h4(e15))) | ~ (h4(op1(e13, e14)) = op2(h4(e13), h4(e14))) | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e15)) = op2(h4(e12), h4(e15))) | ~ (h4(op1(e12, e14)) = op2(h4(e12), h4(e14))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e15)) = op2(h4(e11), h4(e15))) | ~ (h4(op1(e11, e14)) = op2(h4(e11), h4(e14))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e15)) = op2(h4(e10), h4(e15))) | ~ (h4(op1(e10, e14)) = op2(h4(e10), h4(e14))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) & ((~ (e25 = h3(e15)) & ~ (e25 = h3(e14)) & ~ (e25 = h3(e13)) & ~ (e25 = h3(e12)) & ~ (e25 = h3(e11)) & ~ (e25 = h3(e10))) | (~ (e24 = h3(e15)) & ~ (e24 = h3(e14)) & ~ (e24 = h3(e13)) & ~ (e24 = h3(e12)) & ~ (e24 = h3(e11)) & ~ (e24 = h3(e10))) | (~ (e23 = h3(e15)) & ~ (e23 = h3(e14)) & ~ (e23 = h3(e13)) & ~ (e23 = h3(e12)) & ~ (e23 = h3(e11)) & ~ (e23 = h3(e10))) | (~ (e22 = h3(e15)) & ~ (e22 = h3(e14)) & ~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10))) | (~ (e21 = h3(e15)) & ~ (e21 = h3(e14)) & ~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10))) | (~ (e20 = h3(e15)) & ~ (e20 = h3(e14)) & ~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10))) | ~ (h3(op1(e15, e15)) = op2(h3(e15), h3(e15))) | ~ (h3(op1(e15, e14)) = op2(h3(e15), h3(e14))) | ~ (h3(op1(e15, e13)) = op2(h3(e15), h3(e13))) | ~ (h3(op1(e15, e12)) = op2(h3(e15), h3(e12))) | ~ (h3(op1(e15, e11)) = op2(h3(e15), h3(e11))) | ~ (h3(op1(e15, e10)) = op2(h3(e15), h3(e10))) | ~ (h3(op1(e14, e15)) = op2(h3(e14), h3(e15))) | ~ (h3(op1(e14, e14)) = op2(h3(e14), h3(e14))) | ~ (h3(op1(e14, e13)) = op2(h3(e14), h3(e13))) | ~ (h3(op1(e14, e12)) = op2(h3(e14), h3(e12))) | ~ (h3(op1(e14, e11)) = op2(h3(e14), h3(e11))) | ~ (h3(op1(e14, e10)) = op2(h3(e14), h3(e10))) | ~ (h3(op1(e13, e15)) = op2(h3(e13), h3(e15))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), h3(e14))) | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e15)) = op2(h3(e12), h3(e15))) | ~ (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e15)) = op2(h3(e11), h3(e15))) | ~ (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e15)) = op2(h3(e10), h3(e15))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) & ((~ (e25 = h2(e15)) & ~ (e25 = h2(e14)) & ~ (e25 = h2(e13)) & ~ (e25 = h2(e12)) & ~ (e25 = h2(e11)) & ~ (e25 = h2(e10))) | (~ (e24 = h2(e15)) & ~ (e24 = h2(e14)) & ~ (e24 = h2(e13)) & ~ (e24 = h2(e12)) & ~ (e24 = h2(e11)) & ~ (e24 = h2(e10))) | (~ (e23 = h2(e15)) & ~ (e23 = h2(e14)) & ~ (e23 = h2(e13)) & ~ (e23 = h2(e12)) & ~ (e23 = h2(e11)) & ~ (e23 = h2(e10))) | (~ (e22 = h2(e15)) & ~ (e22 = h2(e14)) & ~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10))) | (~ (e21 = h2(e15)) & ~ (e21 = h2(e14)) & ~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10))) | (~ (e20 = h2(e15)) & ~ (e20 = h2(e14)) & ~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10))) | ~ (h2(op1(e15, e15)) = op2(h2(e15), h2(e15))) | ~ (h2(op1(e15, e14)) = op2(h2(e15), h2(e14))) | ~ (h2(op1(e15, e13)) = op2(h2(e15), h2(e13))) | ~ (h2(op1(e15, e12)) = op2(h2(e15), h2(e12))) | ~ (h2(op1(e15, e11)) = op2(h2(e15), h2(e11))) | ~ (h2(op1(e15, e10)) = op2(h2(e15), h2(e10))) | ~ (h2(op1(e14, e15)) = op2(h2(e14), h2(e15))) | ~ (h2(op1(e14, e14)) = op2(h2(e14), h2(e14))) | ~ (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) | ~ (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) | ~ (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) | ~ (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) | ~ (h2(op1(e13, e15)) = op2(h2(e13), h2(e15))) | ~ (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e15)) = op2(h2(e12), h2(e15))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e15)) = op2(h2(e11), h2(e15))) | ~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e15)) = op2(h2(e10), h2(e15))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) & ((~ (e25 = h1(e15)) & ~ (e25 = h1(e14)) & ~ (e25 = h1(e13)) & ~ (e25 = h1(e12)) & ~ (e25 = h1(e11)) & ~ (e25 = h1(e10))) | (~ (e24 = h1(e15)) & ~ (e24 = h1(e14)) & ~ (e24 = h1(e13)) & ~ (e24 = h1(e12)) & ~ (e24 = h1(e11)) & ~ (e24 = h1(e10))) | (~ (e23 = h1(e15)) & ~ (e23 = h1(e14)) & ~ (e23 = h1(e13)) & ~ (e23 = h1(e12)) & ~ (e23 = h1(e11)) & ~ (e23 = h1(e10))) | (~ (e22 = h1(e15)) & ~ (e22 = h1(e14)) & ~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10))) | (~ (e21 = h1(e15)) & ~ (e21 = h1(e14)) & ~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10))) | (~ (e20 = h1(e15)) & ~ (e20 = h1(e14)) & ~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10))) | ~ (h1(op1(e15, e15)) = op2(h1(e15), h1(e15))) | ~ (h1(op1(e15, e14)) = op2(h1(e15), h1(e14))) | ~ (h1(op1(e15, e13)) = op2(h1(e15), h1(e13))) | ~ (h1(op1(e15, e12)) = op2(h1(e15), h1(e12))) | ~ (h1(op1(e15, e11)) = op2(h1(e15), h1(e11))) | ~ (h1(op1(e15, e10)) = op2(h1(e15), h1(e10))) | ~ (h1(op1(e14, e15)) = op2(h1(e14), h1(e15))) | ~ (h1(op1(e14, e14)) = op2(h1(e14), h1(e14))) | ~ (h1(op1(e14, e13)) = op2(h1(e14), h1(e13))) | ~ (h1(op1(e14, e12)) = op2(h1(e14), h1(e12))) | ~ (h1(op1(e14, e11)) = op2(h1(e14), h1(e11))) | ~ (h1(op1(e14, e10)) = op2(h1(e14), h1(e10))) | ~ (h1(op1(e13, e15)) = op2(h1(e13), h1(e15))) | ~ (h1(op1(e13, e14)) = op2(h1(e13), h1(e14))) | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e15)) = op2(h1(e12), h1(e15))) | ~ (h1(op1(e12, e14)) = op2(h1(e12), h1(e14))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e15)) = op2(h1(e11), h1(e15))) | ~ (h1(op1(e11, e14)) = op2(h1(e11), h1(e14))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e15)) = op2(h1(e10), h1(e15))) | ~ (h1(op1(e10, e14)) = op2(h1(e10), h1(e14))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(ennf_transformation, [], [f33])).
fof(f33, plain, ~ ((((e25 = h6(e15)) | (e25 = h6(e14)) | (e25 = h6(e13)) | (e25 = h6(e12)) | (e25 = h6(e11)) | (e25 = h6(e10))) & ((e24 = h6(e15)) | (e24 = h6(e14)) | (e24 = h6(e13)) | (e24 = h6(e12)) | (e24 = h6(e11)) | (e24 = h6(e10))) & ((e23 = h6(e15)) | (e23 = h6(e14)) | (e23 = h6(e13)) | (e23 = h6(e12)) | (e23 = h6(e11)) | (e23 = h6(e10))) & ((e22 = h6(e15)) | (e22 = h6(e14)) | (e22 = h6(e13)) | (e22 = h6(e12)) | (e22 = h6(e11)) | (e22 = h6(e10))) & ((e21 = h6(e15)) | (e21 = h6(e14)) | (e21 = h6(e13)) | (e21 = h6(e12)) | (e21 = h6(e11)) | (e21 = h6(e10))) & ((e20 = h6(e15)) | (e20 = h6(e14)) | (e20 = h6(e13)) | (e20 = h6(e12)) | (e20 = h6(e11)) | (e20 = h6(e10))) & (h6(op1(e15, e15)) = op2(h6(e15), h6(e15))) & (h6(op1(e15, e14)) = op2(h6(e15), h6(e14))) & (h6(op1(e15, e13)) = op2(h6(e15), h6(e13))) & (h6(op1(e15, e12)) = op2(h6(e15), h6(e12))) & (h6(op1(e15, e11)) = op2(h6(e15), h6(e11))) & (h6(op1(e15, e10)) = op2(h6(e15), h6(e10))) & (h6(op1(e14, e15)) = op2(h6(e14), h6(e15))) & (h6(op1(e14, e14)) = op2(h6(e14), h6(e14))) & (h6(op1(e14, e13)) = op2(h6(e14), h6(e13))) & (h6(op1(e14, e12)) = op2(h6(e14), h6(e12))) & (h6(op1(e14, e11)) = op2(h6(e14), h6(e11))) & (h6(op1(e14, e10)) = op2(h6(e14), h6(e10))) & (h6(op1(e13, e15)) = op2(h6(e13), h6(e15))) & (h6(op1(e13, e14)) = op2(h6(e13), h6(e14))) & (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) & (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) & (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) & (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) & (h6(op1(e12, e15)) = op2(h6(e12), h6(e15))) & (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) & (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) & (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) & (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) & (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) & (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) & (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) & (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) & (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) & (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) & (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) & (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) & (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) & (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) & (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) & (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) & (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))) | (((e25 = h5(e15)) | (e25 = h5(e14)) | (e25 = h5(e13)) | (e25 = h5(e12)) | (e25 = h5(e11)) | (e25 = h5(e10))) & ((e24 = h5(e15)) | (e24 = h5(e14)) | (e24 = h5(e13)) | (e24 = h5(e12)) | (e24 = h5(e11)) | (e24 = h5(e10))) & ((e23 = h5(e15)) | (e23 = h5(e14)) | (e23 = h5(e13)) | (e23 = h5(e12)) | (e23 = h5(e11)) | (e23 = h5(e10))) & ((e22 = h5(e15)) | (e22 = h5(e14)) | (e22 = h5(e13)) | (e22 = h5(e12)) | (e22 = h5(e11)) | (e22 = h5(e10))) & ((e21 = h5(e15)) | (e21 = h5(e14)) | (e21 = h5(e13)) | (e21 = h5(e12)) | (e21 = h5(e11)) | (e21 = h5(e10))) & ((e20 = h5(e15)) | (e20 = h5(e14)) | (e20 = h5(e13)) | (e20 = h5(e12)) | (e20 = h5(e11)) | (e20 = h5(e10))) & (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) & (h5(op1(e15, e14)) = op2(h5(e15), h5(e14))) & (h5(op1(e15, e13)) = op2(h5(e15), h5(e13))) & (h5(op1(e15, e12)) = op2(h5(e15), h5(e12))) & (h5(op1(e15, e11)) = op2(h5(e15), h5(e11))) & (h5(op1(e15, e10)) = op2(h5(e15), h5(e10))) & (h5(op1(e14, e15)) = op2(h5(e14), h5(e15))) & (h5(op1(e14, e14)) = op2(h5(e14), h5(e14))) & (h5(op1(e14, e13)) = op2(h5(e14), h5(e13))) & (h5(op1(e14, e12)) = op2(h5(e14), h5(e12))) & (h5(op1(e14, e11)) = op2(h5(e14), h5(e11))) & (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) & (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) & (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) & (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) & (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) & (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) & (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) & (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) & (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) & (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) & (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) & (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) & (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) & (h5(op1(e11, e15)) = op2(h5(e11), h5(e15))) & (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) & (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) & (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) & (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) & (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) & (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) & (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) & (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) & (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) & (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) & (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))) | (((e25 = h4(e15)) | (e25 = h4(e14)) | (e25 = h4(e13)) | (e25 = h4(e12)) | (e25 = h4(e11)) | (e25 = h4(e10))) & ((e24 = h4(e15)) | (e24 = h4(e14)) | (e24 = h4(e13)) | (e24 = h4(e12)) | (e24 = h4(e11)) | (e24 = h4(e10))) & ((e23 = h4(e15)) | (e23 = h4(e14)) | (e23 = h4(e13)) | (e23 = h4(e12)) | (e23 = h4(e11)) | (e23 = h4(e10))) & ((e22 = h4(e15)) | (e22 = h4(e14)) | (e22 = h4(e13)) | (e22 = h4(e12)) | (e22 = h4(e11)) | (e22 = h4(e10))) & ((e21 = h4(e15)) | (e21 = h4(e14)) | (e21 = h4(e13)) | (e21 = h4(e12)) | (e21 = h4(e11)) | (e21 = h4(e10))) & ((e20 = h4(e15)) | (e20 = h4(e14)) | (e20 = h4(e13)) | (e20 = h4(e12)) | (e20 = h4(e11)) | (e20 = h4(e10))) & (h4(op1(e15, e15)) = op2(h4(e15), h4(e15))) & (h4(op1(e15, e14)) = op2(h4(e15), h4(e14))) & (h4(op1(e15, e13)) = op2(h4(e15), h4(e13))) & (h4(op1(e15, e12)) = op2(h4(e15), h4(e12))) & (h4(op1(e15, e11)) = op2(h4(e15), h4(e11))) & (h4(op1(e15, e10)) = op2(h4(e15), h4(e10))) & (h4(op1(e14, e15)) = op2(h4(e14), h4(e15))) & (h4(op1(e14, e14)) = op2(h4(e14), h4(e14))) & (h4(op1(e14, e13)) = op2(h4(e14), h4(e13))) & (h4(op1(e14, e12)) = op2(h4(e14), h4(e12))) & (h4(op1(e14, e11)) = op2(h4(e14), h4(e11))) & (h4(op1(e14, e10)) = op2(h4(e14), h4(e10))) & (h4(op1(e13, e15)) = op2(h4(e13), h4(e15))) & (h4(op1(e13, e14)) = op2(h4(e13), h4(e14))) & (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) & (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) & (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) & (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) & (h4(op1(e12, e15)) = op2(h4(e12), h4(e15))) & (h4(op1(e12, e14)) = op2(h4(e12), h4(e14))) & (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) & (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) & (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) & (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) & (h4(op1(e11, e15)) = op2(h4(e11), h4(e15))) & (h4(op1(e11, e14)) = op2(h4(e11), h4(e14))) & (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) & (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) & (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) & (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) & (h4(op1(e10, e15)) = op2(h4(e10), h4(e15))) & (h4(op1(e10, e14)) = op2(h4(e10), h4(e14))) & (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) & (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) & (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) & (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) | (((e25 = h3(e15)) | (e25 = h3(e14)) | (e25 = h3(e13)) | (e25 = h3(e12)) | (e25 = h3(e11)) | (e25 = h3(e10))) & ((e24 = h3(e15)) | (e24 = h3(e14)) | (e24 = h3(e13)) | (e24 = h3(e12)) | (e24 = h3(e11)) | (e24 = h3(e10))) & ((e23 = h3(e15)) | (e23 = h3(e14)) | (e23 = h3(e13)) | (e23 = h3(e12)) | (e23 = h3(e11)) | (e23 = h3(e10))) & ((e22 = h3(e15)) | (e22 = h3(e14)) | (e22 = h3(e13)) | (e22 = h3(e12)) | (e22 = h3(e11)) | (e22 = h3(e10))) & ((e21 = h3(e15)) | (e21 = h3(e14)) | (e21 = h3(e13)) | (e21 = h3(e12)) | (e21 = h3(e11)) | (e21 = h3(e10))) & ((e20 = h3(e15)) | (e20 = h3(e14)) | (e20 = h3(e13)) | (e20 = h3(e12)) | (e20 = h3(e11)) | (e20 = h3(e10))) & (h3(op1(e15, e15)) = op2(h3(e15), h3(e15))) & (h3(op1(e15, e14)) = op2(h3(e15), h3(e14))) & (h3(op1(e15, e13)) = op2(h3(e15), h3(e13))) & (h3(op1(e15, e12)) = op2(h3(e15), h3(e12))) & (h3(op1(e15, e11)) = op2(h3(e15), h3(e11))) & (h3(op1(e15, e10)) = op2(h3(e15), h3(e10))) & (h3(op1(e14, e15)) = op2(h3(e14), h3(e15))) & (h3(op1(e14, e14)) = op2(h3(e14), h3(e14))) & (h3(op1(e14, e13)) = op2(h3(e14), h3(e13))) & (h3(op1(e14, e12)) = op2(h3(e14), h3(e12))) & (h3(op1(e14, e11)) = op2(h3(e14), h3(e11))) & (h3(op1(e14, e10)) = op2(h3(e14), h3(e10))) & (h3(op1(e13, e15)) = op2(h3(e13), h3(e15))) & (h3(op1(e13, e14)) = op2(h3(e13), h3(e14))) & (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) & (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) & (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) & (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) & (h3(op1(e12, e15)) = op2(h3(e12), h3(e15))) & (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) & (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) & (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) & (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) & (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) & (h3(op1(e11, e15)) = op2(h3(e11), h3(e15))) & (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) & (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) & (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) & (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) & (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) & (h3(op1(e10, e15)) = op2(h3(e10), h3(e15))) & (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) & (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) & (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) & (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) & (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) | (((e25 = h2(e15)) | (e25 = h2(e14)) | (e25 = h2(e13)) | (e25 = h2(e12)) | (e25 = h2(e11)) | (e25 = h2(e10))) & ((e24 = h2(e15)) | (e24 = h2(e14)) | (e24 = h2(e13)) | (e24 = h2(e12)) | (e24 = h2(e11)) | (e24 = h2(e10))) & ((e23 = h2(e15)) | (e23 = h2(e14)) | (e23 = h2(e13)) | (e23 = h2(e12)) | (e23 = h2(e11)) | (e23 = h2(e10))) & ((e22 = h2(e15)) | (e22 = h2(e14)) | (e22 = h2(e13)) | (e22 = h2(e12)) | (e22 = h2(e11)) | (e22 = h2(e10))) & ((e21 = h2(e15)) | (e21 = h2(e14)) | (e21 = h2(e13)) | (e21 = h2(e12)) | (e21 = h2(e11)) | (e21 = h2(e10))) & ((e20 = h2(e15)) | (e20 = h2(e14)) | (e20 = h2(e13)) | (e20 = h2(e12)) | (e20 = h2(e11)) | (e20 = h2(e10))) & (h2(op1(e15, e15)) = op2(h2(e15), h2(e15))) & (h2(op1(e15, e14)) = op2(h2(e15), h2(e14))) & (h2(op1(e15, e13)) = op2(h2(e15), h2(e13))) & (h2(op1(e15, e12)) = op2(h2(e15), h2(e12))) & (h2(op1(e15, e11)) = op2(h2(e15), h2(e11))) & (h2(op1(e15, e10)) = op2(h2(e15), h2(e10))) & (h2(op1(e14, e15)) = op2(h2(e14), h2(e15))) & (h2(op1(e14, e14)) = op2(h2(e14), h2(e14))) & (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) & (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) & (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) & (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) & (h2(op1(e13, e15)) = op2(h2(e13), h2(e15))) & (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) & (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) & (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) & (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) & (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) & (h2(op1(e12, e15)) = op2(h2(e12), h2(e15))) & (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) & (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) & (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) & (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) & (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) & (h2(op1(e11, e15)) = op2(h2(e11), h2(e15))) & (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) & (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) & (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) & (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) & (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) & (h2(op1(e10, e15)) = op2(h2(e10), h2(e15))) & (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) & (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) & (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) & (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) & (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) | (((e25 = h1(e15)) | (e25 = h1(e14)) | (e25 = h1(e13)) | (e25 = h1(e12)) | (e25 = h1(e11)) | (e25 = h1(e10))) & ((e24 = h1(e15)) | (e24 = h1(e14)) | (e24 = h1(e13)) | (e24 = h1(e12)) | (e24 = h1(e11)) | (e24 = h1(e10))) & ((e23 = h1(e15)) | (e23 = h1(e14)) | (e23 = h1(e13)) | (e23 = h1(e12)) | (e23 = h1(e11)) | (e23 = h1(e10))) & ((e22 = h1(e15)) | (e22 = h1(e14)) | (e22 = h1(e13)) | (e22 = h1(e12)) | (e22 = h1(e11)) | (e22 = h1(e10))) & ((e21 = h1(e15)) | (e21 = h1(e14)) | (e21 = h1(e13)) | (e21 = h1(e12)) | (e21 = h1(e11)) | (e21 = h1(e10))) & ((e20 = h1(e15)) | (e20 = h1(e14)) | (e20 = h1(e13)) | (e20 = h1(e12)) | (e20 = h1(e11)) | (e20 = h1(e10))) & (h1(op1(e15, e15)) = op2(h1(e15), h1(e15))) & (h1(op1(e15, e14)) = op2(h1(e15), h1(e14))) & (h1(op1(e15, e13)) = op2(h1(e15), h1(e13))) & (h1(op1(e15, e12)) = op2(h1(e15), h1(e12))) & (h1(op1(e15, e11)) = op2(h1(e15), h1(e11))) & (h1(op1(e15, e10)) = op2(h1(e15), h1(e10))) & (h1(op1(e14, e15)) = op2(h1(e14), h1(e15))) & (h1(op1(e14, e14)) = op2(h1(e14), h1(e14))) & (h1(op1(e14, e13)) = op2(h1(e14), h1(e13))) & (h1(op1(e14, e12)) = op2(h1(e14), h1(e12))) & (h1(op1(e14, e11)) = op2(h1(e14), h1(e11))) & (h1(op1(e14, e10)) = op2(h1(e14), h1(e10))) & (h1(op1(e13, e15)) = op2(h1(e13), h1(e15))) & (h1(op1(e13, e14)) = op2(h1(e13), h1(e14))) & (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) & (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) & (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) & (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) & (h1(op1(e12, e15)) = op2(h1(e12), h1(e15))) & (h1(op1(e12, e14)) = op2(h1(e12), h1(e14))) & (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) & (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) & (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) & (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) & (h1(op1(e11, e15)) = op2(h1(e11), h1(e15))) & (h1(op1(e11, e14)) = op2(h1(e11), h1(e14))) & (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) & (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) & (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) & (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) & (h1(op1(e10, e15)) = op2(h1(e10), h1(e15))) & (h1(op1(e10, e14)) = op2(h1(e10), h1(e14))) & (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) & (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) & (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) & (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(negated_conjecture, [], [f32])).
fof(f32, plain, ~ ((((e25 = h6(e15)) | (e25 = h6(e14)) | (e25 = h6(e13)) | (e25 = h6(e12)) | (e25 = h6(e11)) | (e25 = h6(e10))) & ((e24 = h6(e15)) | (e24 = h6(e14)) | (e24 = h6(e13)) | (e24 = h6(e12)) | (e24 = h6(e11)) | (e24 = h6(e10))) & ((e23 = h6(e15)) | (e23 = h6(e14)) | (e23 = h6(e13)) | (e23 = h6(e12)) | (e23 = h6(e11)) | (e23 = h6(e10))) & ((e22 = h6(e15)) | (e22 = h6(e14)) | (e22 = h6(e13)) | (e22 = h6(e12)) | (e22 = h6(e11)) | (e22 = h6(e10))) & ((e21 = h6(e15)) | (e21 = h6(e14)) | (e21 = h6(e13)) | (e21 = h6(e12)) | (e21 = h6(e11)) | (e21 = h6(e10))) & ((e20 = h6(e15)) | (e20 = h6(e14)) | (e20 = h6(e13)) | (e20 = h6(e12)) | (e20 = h6(e11)) | (e20 = h6(e10))) & (h6(op1(e15, e15)) = op2(h6(e15), h6(e15))) & (h6(op1(e15, e14)) = op2(h6(e15), h6(e14))) & (h6(op1(e15, e13)) = op2(h6(e15), h6(e13))) & (h6(op1(e15, e12)) = op2(h6(e15), h6(e12))) & (h6(op1(e15, e11)) = op2(h6(e15), h6(e11))) & (h6(op1(e15, e10)) = op2(h6(e15), h6(e10))) & (h6(op1(e14, e15)) = op2(h6(e14), h6(e15))) & (h6(op1(e14, e14)) = op2(h6(e14), h6(e14))) & (h6(op1(e14, e13)) = op2(h6(e14), h6(e13))) & (h6(op1(e14, e12)) = op2(h6(e14), h6(e12))) & (h6(op1(e14, e11)) = op2(h6(e14), h6(e11))) & (h6(op1(e14, e10)) = op2(h6(e14), h6(e10))) & (h6(op1(e13, e15)) = op2(h6(e13), h6(e15))) & (h6(op1(e13, e14)) = op2(h6(e13), h6(e14))) & (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) & (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) & (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) & (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) & (h6(op1(e12, e15)) = op2(h6(e12), h6(e15))) & (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) & (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) & (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) & (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) & (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) & (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) & (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) & (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) & (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) & (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) & (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) & (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) & (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) & (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) & (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) & (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) & (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))) | (((e25 = h5(e15)) | (e25 = h5(e14)) | (e25 = h5(e13)) | (e25 = h5(e12)) | (e25 = h5(e11)) | (e25 = h5(e10))) & ((e24 = h5(e15)) | (e24 = h5(e14)) | (e24 = h5(e13)) | (e24 = h5(e12)) | (e24 = h5(e11)) | (e24 = h5(e10))) & ((e23 = h5(e15)) | (e23 = h5(e14)) | (e23 = h5(e13)) | (e23 = h5(e12)) | (e23 = h5(e11)) | (e23 = h5(e10))) & ((e22 = h5(e15)) | (e22 = h5(e14)) | (e22 = h5(e13)) | (e22 = h5(e12)) | (e22 = h5(e11)) | (e22 = h5(e10))) & ((e21 = h5(e15)) | (e21 = h5(e14)) | (e21 = h5(e13)) | (e21 = h5(e12)) | (e21 = h5(e11)) | (e21 = h5(e10))) & ((e20 = h5(e15)) | (e20 = h5(e14)) | (e20 = h5(e13)) | (e20 = h5(e12)) | (e20 = h5(e11)) | (e20 = h5(e10))) & (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) & (h5(op1(e15, e14)) = op2(h5(e15), h5(e14))) & (h5(op1(e15, e13)) = op2(h5(e15), h5(e13))) & (h5(op1(e15, e12)) = op2(h5(e15), h5(e12))) & (h5(op1(e15, e11)) = op2(h5(e15), h5(e11))) & (h5(op1(e15, e10)) = op2(h5(e15), h5(e10))) & (h5(op1(e14, e15)) = op2(h5(e14), h5(e15))) & (h5(op1(e14, e14)) = op2(h5(e14), h5(e14))) & (h5(op1(e14, e13)) = op2(h5(e14), h5(e13))) & (h5(op1(e14, e12)) = op2(h5(e14), h5(e12))) & (h5(op1(e14, e11)) = op2(h5(e14), h5(e11))) & (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) & (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) & (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) & (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) & (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) & (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) & (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) & (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) & (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) & (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) & (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) & (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) & (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) & (h5(op1(e11, e15)) = op2(h5(e11), h5(e15))) & (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) & (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) & (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) & (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) & (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) & (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) & (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) & (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) & (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) & (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) & (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))) | (((e25 = h4(e15)) | (e25 = h4(e14)) | (e25 = h4(e13)) | (e25 = h4(e12)) | (e25 = h4(e11)) | (e25 = h4(e10))) & ((e24 = h4(e15)) | (e24 = h4(e14)) | (e24 = h4(e13)) | (e24 = h4(e12)) | (e24 = h4(e11)) | (e24 = h4(e10))) & ((e23 = h4(e15)) | (e23 = h4(e14)) | (e23 = h4(e13)) | (e23 = h4(e12)) | (e23 = h4(e11)) | (e23 = h4(e10))) & ((e22 = h4(e15)) | (e22 = h4(e14)) | (e22 = h4(e13)) | (e22 = h4(e12)) | (e22 = h4(e11)) | (e22 = h4(e10))) & ((e21 = h4(e15)) | (e21 = h4(e14)) | (e21 = h4(e13)) | (e21 = h4(e12)) | (e21 = h4(e11)) | (e21 = h4(e10))) & ((e20 = h4(e15)) | (e20 = h4(e14)) | (e20 = h4(e13)) | (e20 = h4(e12)) | (e20 = h4(e11)) | (e20 = h4(e10))) & (h4(op1(e15, e15)) = op2(h4(e15), h4(e15))) & (h4(op1(e15, e14)) = op2(h4(e15), h4(e14))) & (h4(op1(e15, e13)) = op2(h4(e15), h4(e13))) & (h4(op1(e15, e12)) = op2(h4(e15), h4(e12))) & (h4(op1(e15, e11)) = op2(h4(e15), h4(e11))) & (h4(op1(e15, e10)) = op2(h4(e15), h4(e10))) & (h4(op1(e14, e15)) = op2(h4(e14), h4(e15))) & (h4(op1(e14, e14)) = op2(h4(e14), h4(e14))) & (h4(op1(e14, e13)) = op2(h4(e14), h4(e13))) & (h4(op1(e14, e12)) = op2(h4(e14), h4(e12))) & (h4(op1(e14, e11)) = op2(h4(e14), h4(e11))) & (h4(op1(e14, e10)) = op2(h4(e14), h4(e10))) & (h4(op1(e13, e15)) = op2(h4(e13), h4(e15))) & (h4(op1(e13, e14)) = op2(h4(e13), h4(e14))) & (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) & (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) & (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) & (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) & (h4(op1(e12, e15)) = op2(h4(e12), h4(e15))) & (h4(op1(e12, e14)) = op2(h4(e12), h4(e14))) & (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) & (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) & (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) & (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) & (h4(op1(e11, e15)) = op2(h4(e11), h4(e15))) & (h4(op1(e11, e14)) = op2(h4(e11), h4(e14))) & (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) & (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) & (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) & (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) & (h4(op1(e10, e15)) = op2(h4(e10), h4(e15))) & (h4(op1(e10, e14)) = op2(h4(e10), h4(e14))) & (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) & (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) & (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) & (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) | (((e25 = h3(e15)) | (e25 = h3(e14)) | (e25 = h3(e13)) | (e25 = h3(e12)) | (e25 = h3(e11)) | (e25 = h3(e10))) & ((e24 = h3(e15)) | (e24 = h3(e14)) | (e24 = h3(e13)) | (e24 = h3(e12)) | (e24 = h3(e11)) | (e24 = h3(e10))) & ((e23 = h3(e15)) | (e23 = h3(e14)) | (e23 = h3(e13)) | (e23 = h3(e12)) | (e23 = h3(e11)) | (e23 = h3(e10))) & ((e22 = h3(e15)) | (e22 = h3(e14)) | (e22 = h3(e13)) | (e22 = h3(e12)) | (e22 = h3(e11)) | (e22 = h3(e10))) & ((e21 = h3(e15)) | (e21 = h3(e14)) | (e21 = h3(e13)) | (e21 = h3(e12)) | (e21 = h3(e11)) | (e21 = h3(e10))) & ((e20 = h3(e15)) | (e20 = h3(e14)) | (e20 = h3(e13)) | (e20 = h3(e12)) | (e20 = h3(e11)) | (e20 = h3(e10))) & (h3(op1(e15, e15)) = op2(h3(e15), h3(e15))) & (h3(op1(e15, e14)) = op2(h3(e15), h3(e14))) & (h3(op1(e15, e13)) = op2(h3(e15), h3(e13))) & (h3(op1(e15, e12)) = op2(h3(e15), h3(e12))) & (h3(op1(e15, e11)) = op2(h3(e15), h3(e11))) & (h3(op1(e15, e10)) = op2(h3(e15), h3(e10))) & (h3(op1(e14, e15)) = op2(h3(e14), h3(e15))) & (h3(op1(e14, e14)) = op2(h3(e14), h3(e14))) & (h3(op1(e14, e13)) = op2(h3(e14), h3(e13))) & (h3(op1(e14, e12)) = op2(h3(e14), h3(e12))) & (h3(op1(e14, e11)) = op2(h3(e14), h3(e11))) & (h3(op1(e14, e10)) = op2(h3(e14), h3(e10))) & (h3(op1(e13, e15)) = op2(h3(e13), h3(e15))) & (h3(op1(e13, e14)) = op2(h3(e13), h3(e14))) & (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) & (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) & (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) & (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) & (h3(op1(e12, e15)) = op2(h3(e12), h3(e15))) & (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) & (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) & (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) & (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) & (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) & (h3(op1(e11, e15)) = op2(h3(e11), h3(e15))) & (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) & (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) & (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) & (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) & (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) & (h3(op1(e10, e15)) = op2(h3(e10), h3(e15))) & (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) & (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) & (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) & (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) & (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) | (((e25 = h2(e15)) | (e25 = h2(e14)) | (e25 = h2(e13)) | (e25 = h2(e12)) | (e25 = h2(e11)) | (e25 = h2(e10))) & ((e24 = h2(e15)) | (e24 = h2(e14)) | (e24 = h2(e13)) | (e24 = h2(e12)) | (e24 = h2(e11)) | (e24 = h2(e10))) & ((e23 = h2(e15)) | (e23 = h2(e14)) | (e23 = h2(e13)) | (e23 = h2(e12)) | (e23 = h2(e11)) | (e23 = h2(e10))) & ((e22 = h2(e15)) | (e22 = h2(e14)) | (e22 = h2(e13)) | (e22 = h2(e12)) | (e22 = h2(e11)) | (e22 = h2(e10))) & ((e21 = h2(e15)) | (e21 = h2(e14)) | (e21 = h2(e13)) | (e21 = h2(e12)) | (e21 = h2(e11)) | (e21 = h2(e10))) & ((e20 = h2(e15)) | (e20 = h2(e14)) | (e20 = h2(e13)) | (e20 = h2(e12)) | (e20 = h2(e11)) | (e20 = h2(e10))) & (h2(op1(e15, e15)) = op2(h2(e15), h2(e15))) & (h2(op1(e15, e14)) = op2(h2(e15), h2(e14))) & (h2(op1(e15, e13)) = op2(h2(e15), h2(e13))) & (h2(op1(e15, e12)) = op2(h2(e15), h2(e12))) & (h2(op1(e15, e11)) = op2(h2(e15), h2(e11))) & (h2(op1(e15, e10)) = op2(h2(e15), h2(e10))) & (h2(op1(e14, e15)) = op2(h2(e14), h2(e15))) & (h2(op1(e14, e14)) = op2(h2(e14), h2(e14))) & (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) & (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) & (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) & (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) & (h2(op1(e13, e15)) = op2(h2(e13), h2(e15))) & (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) & (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) & (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) & (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) & (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) & (h2(op1(e12, e15)) = op2(h2(e12), h2(e15))) & (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) & (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) & (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) & (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) & (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) & (h2(op1(e11, e15)) = op2(h2(e11), h2(e15))) & (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) & (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) & (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) & (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) & (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) & (h2(op1(e10, e15)) = op2(h2(e10), h2(e15))) & (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) & (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) & (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) & (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) & (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) | (((e25 = h1(e15)) | (e25 = h1(e14)) | (e25 = h1(e13)) | (e25 = h1(e12)) | (e25 = h1(e11)) | (e25 = h1(e10))) & ((e24 = h1(e15)) | (e24 = h1(e14)) | (e24 = h1(e13)) | (e24 = h1(e12)) | (e24 = h1(e11)) | (e24 = h1(e10))) & ((e23 = h1(e15)) | (e23 = h1(e14)) | (e23 = h1(e13)) | (e23 = h1(e12)) | (e23 = h1(e11)) | (e23 = h1(e10))) & ((e22 = h1(e15)) | (e22 = h1(e14)) | (e22 = h1(e13)) | (e22 = h1(e12)) | (e22 = h1(e11)) | (e22 = h1(e10))) & ((e21 = h1(e15)) | (e21 = h1(e14)) | (e21 = h1(e13)) | (e21 = h1(e12)) | (e21 = h1(e11)) | (e21 = h1(e10))) & ((e20 = h1(e15)) | (e20 = h1(e14)) | (e20 = h1(e13)) | (e20 = h1(e12)) | (e20 = h1(e11)) | (e20 = h1(e10))) & (h1(op1(e15, e15)) = op2(h1(e15), h1(e15))) & (h1(op1(e15, e14)) = op2(h1(e15), h1(e14))) & (h1(op1(e15, e13)) = op2(h1(e15), h1(e13))) & (h1(op1(e15, e12)) = op2(h1(e15), h1(e12))) & (h1(op1(e15, e11)) = op2(h1(e15), h1(e11))) & (h1(op1(e15, e10)) = op2(h1(e15), h1(e10))) & (h1(op1(e14, e15)) = op2(h1(e14), h1(e15))) & (h1(op1(e14, e14)) = op2(h1(e14), h1(e14))) & (h1(op1(e14, e13)) = op2(h1(e14), h1(e13))) & (h1(op1(e14, e12)) = op2(h1(e14), h1(e12))) & (h1(op1(e14, e11)) = op2(h1(e14), h1(e11))) & (h1(op1(e14, e10)) = op2(h1(e14), h1(e10))) & (h1(op1(e13, e15)) = op2(h1(e13), h1(e15))) & (h1(op1(e13, e14)) = op2(h1(e13), h1(e14))) & (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) & (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) & (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) & (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) & (h1(op1(e12, e15)) = op2(h1(e12), h1(e15))) & (h1(op1(e12, e14)) = op2(h1(e12), h1(e14))) & (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) & (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) & (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) & (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) & (h1(op1(e11, e15)) = op2(h1(e11), h1(e15))) & (h1(op1(e11, e14)) = op2(h1(e11), h1(e14))) & (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) & (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) & (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) & (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) & (h1(op1(e10, e15)) = op2(h1(e10), h1(e15))) & (h1(op1(e10, e14)) = op2(h1(e10), h1(e14))) & (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) & (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) & (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) & (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG026+1.p', co1)).
fof(f5240, plain, ~ spl30_563, inference(avatar_split_clause, [], [f5239, f5219])).
fof(f5239, plain, ~ sP20, inference(subsumption_resolution, [], [f1378, f4945])).
fof(f1378, plain, (~ (e20 = h5(e10)) | ~ sP20), inference(cnf_transformation, [], [f77])).
fof(f77, plain, ((~ (e20 = h5(e15)) & ~ (e20 = h5(e14)) & ~ (e20 = h5(e13)) & ~ (e20 = h5(e12)) & ~ (e20 = h5(e11)) & ~ (e20 = h5(e10))) | ~ sP20), inference(nnf_transformation, [], [f57])).
fof(f5216, plain, ~ spl30_559, inference(avatar_split_clause, [], [f5215, f5196])).
fof(f5215, plain, ~ sP21, inference(subsumption_resolution, [], [f1373, f4941])).
fof(f1373, plain, (~ (e21 = h5(e11)) | ~ sP21), inference(cnf_transformation, [], [f76])).
fof(f76, plain, ((~ (e21 = h5(e15)) & ~ (e21 = h5(e14)) & ~ (e21 = h5(e13)) & ~ (e21 = h5(e12)) & ~ (e21 = h5(e11)) & ~ (e21 = h5(e10))) | ~ sP21), inference(nnf_transformation, [], [f58])).
fof(f5192, plain, (~ spl30_555 | ~ spl30_558), inference(avatar_split_clause, [], [f1368, f5189, f5174])).
fof(f1368, plain, (~ (e22 = h5(e12)) | ~ sP22), inference(cnf_transformation, [], [f75])).
fof(f75, plain, ((~ (e22 = h5(e15)) & ~ (e22 = h5(e14)) & ~ (e22 = h5(e13)) & ~ (e22 = h5(e12)) & ~ (e22 = h5(e11)) & ~ (e22 = h5(e10))) | ~ sP22), inference(nnf_transformation, [], [f59])).
fof(f5165, plain, (~ spl30_551 | ~ spl30_553), inference(avatar_split_clause, [], [f1363, f5162, f5152])).
fof(f1363, plain, (~ (e23 = h5(e13)) | ~ sP23), inference(cnf_transformation, [], [f74])).
fof(f74, plain, ((~ (e23 = h5(e15)) & ~ (e23 = h5(e14)) & ~ (e23 = h5(e13)) & ~ (e23 = h5(e12)) & ~ (e23 = h5(e11)) & ~ (e23 = h5(e10))) | ~ sP23), inference(nnf_transformation, [], [f60])).
fof(f5138, plain, ~ spl30_547, inference(avatar_split_clause, [], [f5137, f5129])).
fof(f5137, plain, ~ sP24, inference(subsumption_resolution, [], [f1358, f1312])).
fof(f1358, plain, (~ (e24 = h5(e14)) | ~ sP24), inference(cnf_transformation, [], [f73])).
fof(f73, plain, ((~ (e24 = h5(e15)) & ~ (e24 = h5(e14)) & ~ (e24 = h5(e13)) & ~ (e24 = h5(e12)) & ~ (e24 = h5(e11)) & ~ (e24 = h5(e10))) | ~ sP24), inference(nnf_transformation, [], [f61])).
fof(f4686, plain, spl30_411, inference(avatar_split_clause, [], [f4685, f3249])).
fof(f4685, plain, (e22 = op2(e21, e24)), inference(backward_demodulation, [], [f4681, f4684])).
fof(f4681, plain, (e22 = op2(op2(e24, op2(e24, e24)), e24)), inference(forward_demodulation, [], [f1285, f589])).
fof(f1285, plain, (e22 = op2(op2(op2(e24, e24), e24), e24)), inference(cnf_transformation, [], [f25])).
fof(f4683, plain, spl30_378, inference(avatar_split_clause, [], [f4682, f3111])).
fof(f4682, plain, (e25 = op2(e22, e24)), inference(backward_demodulation, [], [f4679, f4681])).
fof(f4679, plain, (e25 = op2(op2(op2(e24, op2(e24, e24)), e24), e24)), inference(forward_demodulation, [], [f1287, f589])).
fof(f1287, plain, (e25 = op2(op2(op2(op2(e24, e24), e24), e24), e24)), inference(cnf_transformation, [], [f25])).
fof(f4680, plain, spl30_304, inference(avatar_split_clause, [], [f1286, f2803])).
fof(f1286, plain, (e23 = op2(e24, e24)), inference(cnf_transformation, [], [f25])).
fof(f4677, plain, spl30_153, inference(avatar_split_clause, [], [f4676, f2174])).
fof(f4676, plain, (e12 = op1(e11, e14)), inference(backward_demodulation, [], [f4672, f4675])).
fof(f4672, plain, (e12 = op1(op1(e14, op1(e14, e14)), e14)), inference(forward_demodulation, [], [f1280, f306])).
fof(f1280, plain, (e12 = op1(op1(op1(e14, e14), e14), e14)), inference(cnf_transformation, [], [f24])).
fof(f4674, plain, spl30_120, inference(avatar_split_clause, [], [f4673, f2036])).
fof(f4673, plain, (e15 = op1(e12, e14)), inference(backward_demodulation, [], [f4670, f4672])).
fof(f4670, plain, (e15 = op1(op1(op1(e14, op1(e14, e14)), e14), e14)), inference(forward_demodulation, [], [f1282, f306])).
fof(f1282, plain, (e15 = op1(op1(op1(op1(e14, e14), e14), e14), e14)), inference(cnf_transformation, [], [f24])).
fof(f4671, plain, spl30_46, inference(avatar_split_clause, [], [f1281, f1728])).
fof(f1281, plain, (e13 = op1(e14, e14)), inference(cnf_transformation, [], [f24])).
fof(f3697, plain, (~ spl30_231 | spl30_245), inference(avatar_split_clause, [], [f697, f2557, f2499])).
fof(f2499, plain, (spl30_231 <=> (e12 = inv1(e14))), introduced(avatar_definition, [new_symbols(naming, [spl30_231])])).
fof(f697, plain, ((e14 = inv1(e12)) | ~ (e12 = inv1(e14))), inference(cnf_transformation, [], [f34])).
fof(f34, plain, (((e15 = inv1(e15)) | ~ (e15 = inv1(e15))) & ((e15 = inv1(e14)) | ~ (e14 = inv1(e15))) & ((e15 = inv1(e13)) | ~ (e13 = inv1(e15))) & ((e15 = inv1(e12)) | ~ (e12 = inv1(e15))) & ((e15 = inv1(e11)) | ~ (e11 = inv1(e15))) & ((e15 = inv1(e10)) | ~ (e10 = inv1(e15))) & ((e14 = inv1(e15)) | ~ (e15 = inv1(e14))) & ((e14 = inv1(e14)) | ~ (e14 = inv1(e14))) & ((e14 = inv1(e13)) | ~ (e13 = inv1(e14))) & ((e14 = inv1(e12)) | ~ (e12 = inv1(e14))) & ((e14 = inv1(e11)) | ~ (e11 = inv1(e14))) & ((e14 = inv1(e10)) | ~ (e10 = inv1(e14))) & ((e13 = inv1(e15)) | ~ (e15 = inv1(e13))) & ((e13 = inv1(e14)) | ~ (e14 = inv1(e13))) & ((e13 = inv1(e13)) | ~ (e13 = inv1(e13))) & ((e13 = inv1(e12)) | ~ (e12 = inv1(e13))) & ((e13 = inv1(e11)) | ~ (e11 = inv1(e13))) & ((e13 = inv1(e10)) | ~ (e10 = inv1(e13))) & ((e12 = inv1(e15)) | ~ (e15 = inv1(e12))) & ((e12 = inv1(e14)) | ~ (e14 = inv1(e12))) & ((e12 = inv1(e13)) | ~ (e13 = inv1(e12))) & ((e12 = inv1(e12)) | ~ (e12 = inv1(e12))) & ((e12 = inv1(e11)) | ~ (e11 = inv1(e12))) & ((e12 = inv1(e10)) | ~ (e10 = inv1(e12))) & ((e11 = inv1(e15)) | ~ (e15 = inv1(e11))) & ((e11 = inv1(e14)) | ~ (e14 = inv1(e11))) & ((e11 = inv1(e13)) | ~ (e13 = inv1(e11))) & ((e11 = inv1(e12)) | ~ (e12 = inv1(e11))) & ((e11 = inv1(e11)) | ~ (e11 = inv1(e11))) & ((e11 = inv1(e10)) | ~ (e10 = inv1(e11))) & ((e10 = inv1(e15)) | ~ (e15 = inv1(e10))) & ((e10 = inv1(e14)) | ~ (e14 = inv1(e10))) & ((e10 = inv1(e13)) | ~ (e13 = inv1(e10))) & ((e10 = inv1(e12)) | ~ (e12 = inv1(e10))) & ((e10 = inv1(e11)) | ~ (e11 = inv1(e10))) & ((e10 = inv1(e10)) | ~ (e10 = inv1(e10)))), inference(ennf_transformation, [], [f11])).
fof(f11, plain, (((e15 = inv1(e15)) => (e15 = inv1(e15))) & ((e14 = inv1(e15)) => (e15 = inv1(e14))) & ((e13 = inv1(e15)) => (e15 = inv1(e13))) & ((e12 = inv1(e15)) => (e15 = inv1(e12))) & ((e11 = inv1(e15)) => (e15 = inv1(e11))) & ((e10 = inv1(e15)) => (e15 = inv1(e10))) & ((e15 = inv1(e14)) => (e14 = inv1(e15))) & ((e14 = inv1(e14)) => (e14 = inv1(e14))) & ((e13 = inv1(e14)) => (e14 = inv1(e13))) & ((e12 = inv1(e14)) => (e14 = inv1(e12))) & ((e11 = inv1(e14)) => (e14 = inv1(e11))) & ((e10 = inv1(e14)) => (e14 = inv1(e10))) & ((e15 = inv1(e13)) => (e13 = inv1(e15))) & ((e14 = inv1(e13)) => (e13 = inv1(e14))) & ((e13 = inv1(e13)) => (e13 = inv1(e13))) & ((e12 = inv1(e13)) => (e13 = inv1(e12))) & ((e11 = inv1(e13)) => (e13 = inv1(e11))) & ((e10 = inv1(e13)) => (e13 = inv1(e10))) & ((e15 = inv1(e12)) => (e12 = inv1(e15))) & ((e14 = inv1(e12)) => (e12 = inv1(e14))) & ((e13 = inv1(e12)) => (e12 = inv1(e13))) & ((e12 = inv1(e12)) => (e12 = inv1(e12))) & ((e11 = inv1(e12)) => (e12 = inv1(e11))) & ((e10 = inv1(e12)) => (e12 = inv1(e10))) & ((e15 = inv1(e11)) => (e11 = inv1(e15))) & ((e14 = inv1(e11)) => (e11 = inv1(e14))) & ((e13 = inv1(e11)) => (e11 = inv1(e13))) & ((e12 = inv1(e11)) => (e11 = inv1(e12))) & ((e11 = inv1(e11)) => (e11 = inv1(e11))) & ((e10 = inv1(e11)) => (e11 = inv1(e10))) & ((e15 = inv1(e10)) => (e10 = inv1(e15))) & ((e14 = inv1(e10)) => (e10 = inv1(e14))) & ((e13 = inv1(e10)) => (e10 = inv1(e13))) & ((e12 = inv1(e10)) => (e10 = inv1(e12))) & ((e11 = inv1(e10)) => (e10 = inv1(e11))) & ((e10 = inv1(e10)) => (e10 = inv1(e10)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG026+1.p', ax11)).
fof(f3539, plain, (spl30_475 | spl30_476 | spl30_477 | spl30_478 | spl30_479 | spl30_480), inference(avatar_split_clause, [], [f645, f3536, f3532, f3528, f3524, f3520, f3516])).
fof(f645, plain, ((e25 = unit2) | (e24 = unit2) | (e23 = unit2) | (e22 = unit2) | (e21 = unit2) | (e20 = unit2)), inference(cnf_transformation, [], [f7])).
fof(f2514, plain, (spl30_229 | spl30_230 | spl30_231 | spl30_232 | spl30_233 | spl30_234), inference(avatar_split_clause, [], [f379, f2511, f2507, f2503, f2499, f2495, f2491])).
fof(f379, plain, ((e15 = inv1(e14)) | (e14 = inv1(e14)) | (e13 = inv1(e14)) | (e12 = inv1(e14)) | (e11 = inv1(e14)) | (e10 = inv1(e14))), inference(cnf_transformation, [], [f4])).
fof(f2339, plain, (spl30_187 | spl30_188 | spl30_189 | spl30_190 | spl30_191 | spl30_192), inference(avatar_split_clause, [], [f102, f2336, f2332, f2328, f2324, f2320, f2316])).
fof(f102, plain, ((e15 = op1(e10, e14)) | (e14 = op1(e10, e14)) | (e13 = op1(e10, e14)) | (e12 = op1(e10, e14)) | (e11 = op1(e10, e14)) | (e10 = op1(e10, e14))), inference(cnf_transformation, [], [f1])).
fof(f1, plain, (((e15 = op1(e15, e15)) | (e14 = op1(e15, e15)) | (e13 = op1(e15, e15)) | (e12 = op1(e15, e15)) | (e11 = op1(e15, e15)) | (e10 = op1(e15, e15))) & ((e15 = op1(e15, e14)) | (e14 = op1(e15, e14)) | (e13 = op1(e15, e14)) | (e12 = op1(e15, e14)) | (e11 = op1(e15, e14)) | (e10 = op1(e15, e14))) & ((e15 = op1(e15, e13)) | (e14 = op1(e15, e13)) | (e13 = op1(e15, e13)) | (e12 = op1(e15, e13)) | (e11 = op1(e15, e13)) | (e10 = op1(e15, e13))) & ((e15 = op1(e15, e12)) | (e14 = op1(e15, e12)) | (e13 = op1(e15, e12)) | (e12 = op1(e15, e12)) | (e11 = op1(e15, e12)) | (e10 = op1(e15, e12))) & ((e15 = op1(e15, e11)) | (e14 = op1(e15, e11)) | (e13 = op1(e15, e11)) | (e12 = op1(e15, e11)) | (e11 = op1(e15, e11)) | (e10 = op1(e15, e11))) & ((e15 = op1(e15, e10)) | (e14 = op1(e15, e10)) | (e13 = op1(e15, e10)) | (e12 = op1(e15, e10)) | (e11 = op1(e15, e10)) | (e10 = op1(e15, e10))) & ((e15 = op1(e14, e15)) | (e14 = op1(e14, e15)) | (e13 = op1(e14, e15)) | (e12 = op1(e14, e15)) | (e11 = op1(e14, e15)) | (e10 = op1(e14, e15))) & ((e15 = op1(e14, e14)) | (e14 = op1(e14, e14)) | (e13 = op1(e14, e14)) | (e12 = op1(e14, e14)) | (e11 = op1(e14, e14)) | (e10 = op1(e14, e14))) & ((e15 = op1(e14, e13)) | (e14 = op1(e14, e13)) | (e13 = op1(e14, e13)) | (e12 = op1(e14, e13)) | (e11 = op1(e14, e13)) | (e10 = op1(e14, e13))) & ((e15 = op1(e14, e12)) | (e14 = op1(e14, e12)) | (e13 = op1(e14, e12)) | (e12 = op1(e14, e12)) | (e11 = op1(e14, e12)) | (e10 = op1(e14, e12))) & ((e15 = op1(e14, e11)) | (e14 = op1(e14, e11)) | (e13 = op1(e14, e11)) | (e12 = op1(e14, e11)) | (e11 = op1(e14, e11)) | (e10 = op1(e14, e11))) & ((e15 = op1(e14, e10)) | (e14 = op1(e14, e10)) | (e13 = op1(e14, e10)) | (e12 = op1(e14, e10)) | (e11 = op1(e14, e10)) | (e10 = op1(e14, e10))) & ((e15 = op1(e13, e15)) | (e14 = op1(e13, e15)) | (e13 = op1(e13, e15)) | (e12 = op1(e13, e15)) | (e11 = op1(e13, e15)) | (e10 = op1(e13, e15))) & ((e15 = op1(e13, e14)) | (e14 = op1(e13, e14)) | (e13 = op1(e13, e14)) | (e12 = op1(e13, e14)) | (e11 = op1(e13, e14)) | (e10 = op1(e13, e14))) & ((e15 = op1(e13, e13)) | (e14 = op1(e13, e13)) | (e13 = op1(e13, e13)) | (e12 = op1(e13, e13)) | (e11 = op1(e13, e13)) | (e10 = op1(e13, e13))) & ((e15 = op1(e13, e12)) | (e14 = op1(e13, e12)) | (e13 = op1(e13, e12)) | (e12 = op1(e13, e12)) | (e11 = op1(e13, e12)) | (e10 = op1(e13, e12))) & ((e15 = op1(e13, e11)) | (e14 = op1(e13, e11)) | (e13 = op1(e13, e11)) | (e12 = op1(e13, e11)) | (e11 = op1(e13, e11)) | (e10 = op1(e13, e11))) & ((e15 = op1(e13, e10)) | (e14 = op1(e13, e10)) | (e13 = op1(e13, e10)) | (e12 = op1(e13, e10)) | (e11 = op1(e13, e10)) | (e10 = op1(e13, e10))) & ((e15 = op1(e12, e15)) | (e14 = op1(e12, e15)) | (e13 = op1(e12, e15)) | (e12 = op1(e12, e15)) | (e11 = op1(e12, e15)) | (e10 = op1(e12, e15))) & ((e15 = op1(e12, e14)) | (e14 = op1(e12, e14)) | (e13 = op1(e12, e14)) | (e12 = op1(e12, e14)) | (e11 = op1(e12, e14)) | (e10 = op1(e12, e14))) & ((e15 = op1(e12, e13)) | (e14 = op1(e12, e13)) | (e13 = op1(e12, e13)) | (e12 = op1(e12, e13)) | (e11 = op1(e12, e13)) | (e10 = op1(e12, e13))) & ((e15 = op1(e12, e12)) | (e14 = op1(e12, e12)) | (e13 = op1(e12, e12)) | (e12 = op1(e12, e12)) | (e11 = op1(e12, e12)) | (e10 = op1(e12, e12))) & ((e15 = op1(e12, e11)) | (e14 = op1(e12, e11)) | (e13 = op1(e12, e11)) | (e12 = op1(e12, e11)) | (e11 = op1(e12, e11)) | (e10 = op1(e12, e11))) & ((e15 = op1(e12, e10)) | (e14 = op1(e12, e10)) | (e13 = op1(e12, e10)) | (e12 = op1(e12, e10)) | (e11 = op1(e12, e10)) | (e10 = op1(e12, e10))) & ((e15 = op1(e11, e15)) | (e14 = op1(e11, e15)) | (e13 = op1(e11, e15)) | (e12 = op1(e11, e15)) | (e11 = op1(e11, e15)) | (e10 = op1(e11, e15))) & ((e15 = op1(e11, e14)) | (e14 = op1(e11, e14)) | (e13 = op1(e11, e14)) | (e12 = op1(e11, e14)) | (e11 = op1(e11, e14)) | (e10 = op1(e11, e14))) & ((e15 = op1(e11, e13)) | (e14 = op1(e11, e13)) | (e13 = op1(e11, e13)) | (e12 = op1(e11, e13)) | (e11 = op1(e11, e13)) | (e10 = op1(e11, e13))) & ((e15 = op1(e11, e12)) | (e14 = op1(e11, e12)) | (e13 = op1(e11, e12)) | (e12 = op1(e11, e12)) | (e11 = op1(e11, e12)) | (e10 = op1(e11, e12))) & ((e15 = op1(e11, e11)) | (e14 = op1(e11, e11)) | (e13 = op1(e11, e11)) | (e12 = op1(e11, e11)) | (e11 = op1(e11, e11)) | (e10 = op1(e11, e11))) & ((e15 = op1(e11, e10)) | (e14 = op1(e11, e10)) | (e13 = op1(e11, e10)) | (e12 = op1(e11, e10)) | (e11 = op1(e11, e10)) | (e10 = op1(e11, e10))) & ((e15 = op1(e10, e15)) | (e14 = op1(e10, e15)) | (e13 = op1(e10, e15)) | (e12 = op1(e10, e15)) | (e11 = op1(e10, e15)) | (e10 = op1(e10, e15))) & ((e15 = op1(e10, e14)) | (e14 = op1(e10, e14)) | (e13 = op1(e10, e14)) | (e12 = op1(e10, e14)) | (e11 = op1(e10, e14)) | (e10 = op1(e10, e14))) & ((e15 = op1(e10, e13)) | (e14 = op1(e10, e13)) | (e13 = op1(e10, e13)) | (e12 = op1(e10, e13)) | (e11 = op1(e10, e13)) | (e10 = op1(e10, e13))) & ((e15 = op1(e10, e12)) | (e14 = op1(e10, e12)) | (e13 = op1(e10, e12)) | (e12 = op1(e10, e12)) | (e11 = op1(e10, e12)) | (e10 = op1(e10, e12))) & ((e15 = op1(e10, e11)) | (e14 = op1(e10, e11)) | (e13 = op1(e10, e11)) | (e12 = op1(e10, e11)) | (e11 = op1(e10, e11)) | (e10 = op1(e10, e11))) & ((op1(e10, e10) = e15) | (op1(e10, e10) = e14) | (op1(e10, e10) = e13) | (op1(e10, e10) = e12) | (op1(e10, e10) = e11) | (e10 = op1(e10, e10)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG026+1.p', ax1)).