fof(f1532, plain, $false, inference(subsumption_resolution, [], [f1531, f1084])).
fof(f1084, plain, ~ sP0, inference(subsumption_resolution, [], [f1083, f609])).
fof(f609, plain, ~ (e4 = unit), inference(definition_unfolding, [], [f171, f203])).
fof(f203, plain, (e0 = unit), inference(cnf_transformation, [], [f3])).
fof(f3, plain, (e0 = unit), file('/home/ubuntu/library/tptp/Problems/ALG/ALG092+1.p', ax3)).
fof(f171, plain, ~ (e0 = e4), inference(cnf_transformation, [], [f1])).
fof(f1, plain, (~ (e3 = e4) & ~ (e2 = e4) & ~ (e2 = e3) & ~ (e1 = e4) & ~ (e1 = e3) & ~ (e1 = e2) & ~ (e0 = e4) & ~ (e0 = e3) & ~ (e0 = e2) & ~ (e0 = e1)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG092+1.p', ax1)).
fof(f1083, plain, ((e4 = unit) | ~ sP0), inference(forward_demodulation, [], [f1082, f196])).
fof(f196, plain, (e4 = op(e3, e3)), inference(cnf_transformation, [], [f2])).
fof(f2, plain, ((e0 = op(e4, e4)) & (e2 = op(e4, e3)) & (e1 = op(e4, e2)) & (e3 = op(e4, e1)) & (e4 = op(e4, e0)) & (e1 = op(e3, e4)) & (e4 = op(e3, e3)) & (e0 = op(e3, e2)) & (e2 = op(e3, e1)) & (e3 = op(e3, e0)) & (e3 = op(e2, e4)) & (e1 = op(e2, e3)) & (e4 = op(e2, e2)) & (e0 = op(e2, e1)) & (e2 = op(e2, e0)) & (e2 = op(e1, e4)) & (e0 = op(e1, e3)) & (e3 = op(e1, e2)) & (e4 = op(e1, e1)) & (e1 = op(e1, e0)) & (e4 = op(e0, e4)) & (e3 = op(e0, e3)) & (e2 = op(e0, e2)) & (e1 = op(e0, e1)) & (e0 = op(e0, e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG092+1.p', ax2)).
fof(f1082, plain, ((op(e3, e3) = unit) | ~ sP0), inference(forward_demodulation, [], [f1081, f616])).
fof(f616, plain, (e3 = op(e3, unit)), inference(definition_unfolding, [], [f193, f203])).
fof(f193, plain, (e3 = op(e3, e0)), inference(cnf_transformation, [], [f2])).
fof(f1081, plain, ((unit = op(e3, op(e3, unit))) | ~ sP0), inference(subsumption_resolution, [], [f1080, f622])).
fof(f622, plain, (e3 = op(unit, e3)), inference(definition_unfolding, [], [f181, f203])).
fof(f181, plain, (e3 = op(e0, e3)), inference(cnf_transformation, [], [f2])).
fof(f1080, plain, (~ (e3 = op(unit, e3)) | (unit = op(e3, op(e3, unit))) | ~ sP0), inference(forward_demodulation, [], [f814, f616])).
fof(f814, plain, (~ (e3 = op(unit, op(e3, unit))) | (unit = op(e3, op(e3, unit))) | ~ sP0), inference(definition_unfolding, [], [f602, f203, f203, f203, f203])).
fof(f602, plain, (~ (e3 = op(e0, op(e3, e0))) | (e0 = op(e3, op(e3, e0))) | ~ sP0), inference(cnf_transformation, [], [f167])).
fof(f167, plain, (((~ (e4 = op(e0, op(e4, e0))) | (e0 = op(e4, op(e4, e0)))) & (~ (e3 = op(e0, op(e3, e0))) | (e0 = op(e3, op(e3, e0)))) & (~ (e2 = op(e0, op(e2, e0))) | (e0 = op(e2, op(e2, e0)))) & (~ (e1 = op(e0, op(e1, e0))) | (e0 = op(e1, op(e1, e0)))) & (~ (e0 = op(e0, op(e0, e0))) | (e0 = op(e0, op(e0, e0))))) | ~ sP0), inference(nnf_transformation, [], [f7])).
fof(f7, plain, (((~ (e4 = op(e0, op(e4, e0))) | (e0 = op(e4, op(e4, e0)))) & (~ (e3 = op(e0, op(e3, e0))) | (e0 = op(e3, op(e3, e0)))) & (~ (e2 = op(e0, op(e2, e0))) | (e0 = op(e2, op(e2, e0)))) & (~ (e1 = op(e0, op(e1, e0))) | (e0 = op(e1, op(e1, e0)))) & (~ (e0 = op(e0, op(e0, e0))) | (e0 = op(e0, op(e0, e0))))) | ~ sP0), inference(usedef, [], [e7])).
fof(e7, plain, (sP0 <=> ((~ (e4 = op(e0, op(e4, e0))) | (e0 = op(e4, op(e4, e0)))) & (~ (e3 = op(e0, op(e3, e0))) | (e0 = op(e3, op(e3, e0)))) & (~ (e2 = op(e0, op(e2, e0))) | (e0 = op(e2, op(e2, e0)))) & (~ (e1 = op(e0, op(e1, e0))) | (e0 = op(e1, op(e1, e0)))) & (~ (e0 = op(e0, op(e0, e0))) | (e0 = op(e0, op(e0, e0)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f1531, plain, sP0, inference(subsumption_resolution, [], [f1530, f1076])).
fof(f1076, plain, ~ sP1, inference(subsumption_resolution, [], [f1075, f612])).
fof(f612, plain, ~ (e1 = unit), inference(definition_unfolding, [], [f168, f203])).
fof(f168, plain, ~ (e0 = e1), inference(cnf_transformation, [], [f1])).
fof(f1075, plain, ((e1 = unit) | ~ sP1), inference(forward_demodulation, [], [f1074, f615])).
fof(f615, plain, (op(e3, e2) = unit), inference(definition_unfolding, [], [f195, f203])).
fof(f195, plain, (e0 = op(e3, e2)), inference(cnf_transformation, [], [f2])).
fof(f1074, plain, ((e1 = op(e3, e2)) | ~ sP1), inference(forward_demodulation, [], [f1073, f194])).
fof(f194, plain, (e2 = op(e3, e1)), inference(cnf_transformation, [], [f2])).
fof(f1073, plain, ((e1 = op(e3, op(e3, e1))) | ~ sP1), inference(subsumption_resolution, [], [f1072, f185])).
fof(f185, plain, (e3 = op(e1, e2)), inference(cnf_transformation, [], [f2])).
fof(f1072, plain, (~ (e3 = op(e1, e2)) | (e1 = op(e3, op(e3, e1))) | ~ sP1), inference(forward_demodulation, [], [f597, f194])).
fof(f597, plain, (~ (e3 = op(e1, op(e3, e1))) | (e1 = op(e3, op(e3, e1))) | ~ sP1), inference(cnf_transformation, [], [f166])).
fof(f166, plain, (((~ (e4 = op(e1, op(e4, e1))) | (e1 = op(e4, op(e4, e1)))) & (~ (e3 = op(e1, op(e3, e1))) | (e1 = op(e3, op(e3, e1)))) & (~ (e2 = op(e1, op(e2, e1))) | (e1 = op(e2, op(e2, e1)))) & (~ (e1 = op(e1, op(e1, e1))) | (e1 = op(e1, op(e1, e1)))) & (~ (e0 = op(e1, op(e0, e1))) | (e1 = op(e0, op(e0, e1))))) | ~ sP1), inference(nnf_transformation, [], [f8])).
fof(f8, plain, (((~ (e4 = op(e1, op(e4, e1))) | (e1 = op(e4, op(e4, e1)))) & (~ (e3 = op(e1, op(e3, e1))) | (e1 = op(e3, op(e3, e1)))) & (~ (e2 = op(e1, op(e2, e1))) | (e1 = op(e2, op(e2, e1)))) & (~ (e1 = op(e1, op(e1, e1))) | (e1 = op(e1, op(e1, e1)))) & (~ (e0 = op(e1, op(e0, e1))) | (e1 = op(e0, op(e0, e1))))) | ~ sP1), inference(usedef, [], [e8])).
fof(e8, plain, (sP1 <=> ((~ (e4 = op(e1, op(e4, e1))) | (e1 = op(e4, op(e4, e1)))) & (~ (e3 = op(e1, op(e3, e1))) | (e1 = op(e3, op(e3, e1)))) & (~ (e2 = op(e1, op(e2, e1))) | (e1 = op(e2, op(e2, e1)))) & (~ (e1 = op(e1, op(e1, e1))) | (e1 = op(e1, op(e1, e1)))) & (~ (e0 = op(e1, op(e0, e1))) | (e1 = op(e0, op(e0, e1)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f1530, plain, (sP1 | sP0), inference(subsumption_resolution, [], [f1529, f1069])).
fof(f1069, plain, ~ sP2, inference(subsumption_resolution, [], [f1068, f611])).
fof(f611, plain, ~ (e2 = unit), inference(definition_unfolding, [], [f169, f203])).
fof(f169, plain, ~ (e0 = e2), inference(cnf_transformation, [], [f1])).
fof(f1068, plain, ((e2 = unit) | ~ sP2), inference(forward_demodulation, [], [f1067, f619])).
fof(f619, plain, (op(e1, e3) = unit), inference(definition_unfolding, [], [f186, f203])).
fof(f186, plain, (e0 = op(e1, e3)), inference(cnf_transformation, [], [f2])).
fof(f1067, plain, ((e2 = op(e1, e3)) | ~ sP2), inference(forward_demodulation, [], [f1066, f185])).
fof(f1066, plain, ((e2 = op(e1, op(e1, e2))) | ~ sP2), inference(subsumption_resolution, [], [f1065, f191])).
fof(f191, plain, (e1 = op(e2, e3)), inference(cnf_transformation, [], [f2])).
fof(f1065, plain, (~ (e1 = op(e2, e3)) | (e2 = op(e1, op(e1, e2))) | ~ sP2), inference(forward_demodulation, [], [f590, f185])).
fof(f590, plain, (~ (e1 = op(e2, op(e1, e2))) | (e2 = op(e1, op(e1, e2))) | ~ sP2), inference(cnf_transformation, [], [f165])).
fof(f165, plain, (((~ (e4 = op(e2, op(e4, e2))) | (e2 = op(e4, op(e4, e2)))) & (~ (e3 = op(e2, op(e3, e2))) | (e2 = op(e3, op(e3, e2)))) & (~ (e2 = op(e2, op(e2, e2))) | (e2 = op(e2, op(e2, e2)))) & (~ (e1 = op(e2, op(e1, e2))) | (e2 = op(e1, op(e1, e2)))) & (~ (e0 = op(e2, op(e0, e2))) | (e2 = op(e0, op(e0, e2))))) | ~ sP2), inference(nnf_transformation, [], [f9])).
fof(f9, plain, (((~ (e4 = op(e2, op(e4, e2))) | (e2 = op(e4, op(e4, e2)))) & (~ (e3 = op(e2, op(e3, e2))) | (e2 = op(e3, op(e3, e2)))) & (~ (e2 = op(e2, op(e2, e2))) | (e2 = op(e2, op(e2, e2)))) & (~ (e1 = op(e2, op(e1, e2))) | (e2 = op(e1, op(e1, e2)))) & (~ (e0 = op(e2, op(e0, e2))) | (e2 = op(e0, op(e0, e2))))) | ~ sP2), inference(usedef, [], [e9])).
fof(e9, plain, (sP2 <=> ((~ (e4 = op(e2, op(e4, e2))) | (e2 = op(e4, op(e4, e2)))) & (~ (e3 = op(e2, op(e3, e2))) | (e2 = op(e3, op(e3, e2)))) & (~ (e2 = op(e2, op(e2, e2))) | (e2 = op(e2, op(e2, e2)))) & (~ (e1 = op(e2, op(e1, e2))) | (e2 = op(e1, op(e1, e2)))) & (~ (e0 = op(e2, op(e0, e2))) | (e2 = op(e0, op(e0, e2)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f1529, plain, (sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1528, f1060])).
fof(f1060, plain, ~ sP3, inference(subsumption_resolution, [], [f1059, f610])).
fof(f610, plain, ~ (e3 = unit), inference(definition_unfolding, [], [f170, f203])).
fof(f170, plain, ~ (e0 = e3), inference(cnf_transformation, [], [f1])).
fof(f1059, plain, ((e3 = unit) | ~ sP3), inference(forward_demodulation, [], [f1058, f617])).
fof(f617, plain, (op(e2, e1) = unit), inference(definition_unfolding, [], [f189, f203])).
fof(f189, plain, (e0 = op(e2, e1)), inference(cnf_transformation, [], [f2])).
fof(f1058, plain, ((e3 = op(e2, e1)) | ~ sP3), inference(forward_demodulation, [], [f1057, f191])).
fof(f1057, plain, ((e3 = op(e2, op(e2, e3))) | ~ sP3), inference(subsumption_resolution, [], [f1056, f194])).
fof(f1056, plain, (~ (e2 = op(e3, e1)) | (e3 = op(e2, op(e2, e3))) | ~ sP3), inference(forward_demodulation, [], [f586, f191])).
fof(f586, plain, (~ (e2 = op(e3, op(e2, e3))) | (e3 = op(e2, op(e2, e3))) | ~ sP3), inference(cnf_transformation, [], [f164])).
fof(f164, plain, (((~ (e4 = op(e3, op(e4, e3))) | (e3 = op(e4, op(e4, e3)))) & (~ (e3 = op(e3, op(e3, e3))) | (e3 = op(e3, op(e3, e3)))) & (~ (e2 = op(e3, op(e2, e3))) | (e3 = op(e2, op(e2, e3)))) & (~ (e1 = op(e3, op(e1, e3))) | (e3 = op(e1, op(e1, e3)))) & (~ (e0 = op(e3, op(e0, e3))) | (e3 = op(e0, op(e0, e3))))) | ~ sP3), inference(nnf_transformation, [], [f10])).
fof(f10, plain, (((~ (e4 = op(e3, op(e4, e3))) | (e3 = op(e4, op(e4, e3)))) & (~ (e3 = op(e3, op(e3, e3))) | (e3 = op(e3, op(e3, e3)))) & (~ (e2 = op(e3, op(e2, e3))) | (e3 = op(e2, op(e2, e3)))) & (~ (e1 = op(e3, op(e1, e3))) | (e3 = op(e1, op(e1, e3)))) & (~ (e0 = op(e3, op(e0, e3))) | (e3 = op(e0, op(e0, e3))))) | ~ sP3), inference(usedef, [], [e10])).
fof(e10, plain, (sP3 <=> ((~ (e4 = op(e3, op(e4, e3))) | (e3 = op(e4, op(e4, e3)))) & (~ (e3 = op(e3, op(e3, e3))) | (e3 = op(e3, op(e3, e3)))) & (~ (e2 = op(e3, op(e2, e3))) | (e3 = op(e2, op(e2, e3)))) & (~ (e1 = op(e3, op(e1, e3))) | (e3 = op(e1, op(e1, e3)))) & (~ (e0 = op(e3, op(e0, e3))) | (e3 = op(e0, op(e0, e3)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f1528, plain, (sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1527, f1053])).
fof(f1053, plain, ~ sP4, inference(subsumption_resolution, [], [f1052, f176])).
fof(f176, plain, ~ (e2 = e4), inference(cnf_transformation, [], [f1])).
fof(f1052, plain, ((e2 = e4) | ~ sP4), inference(forward_demodulation, [], [f1051, f194])).
fof(f1051, plain, ((e4 = op(e3, e1)) | ~ sP4), inference(forward_demodulation, [], [f1050, f197])).
fof(f197, plain, (e1 = op(e3, e4)), inference(cnf_transformation, [], [f2])).
fof(f1050, plain, ((e4 = op(e3, op(e3, e4))) | ~ sP4), inference(subsumption_resolution, [], [f1049, f199])).
fof(f199, plain, (e3 = op(e4, e1)), inference(cnf_transformation, [], [f2])).
fof(f1049, plain, (~ (e3 = op(e4, e1)) | (e4 = op(e3, op(e3, e4))) | ~ sP4), inference(forward_demodulation, [], [f582, f197])).
fof(f582, plain, (~ (e3 = op(e4, op(e3, e4))) | (e4 = op(e3, op(e3, e4))) | ~ sP4), inference(cnf_transformation, [], [f163])).
fof(f163, plain, (((~ (e4 = op(e4, op(e4, e4))) | (e4 = op(e4, op(e4, e4)))) & (~ (e3 = op(e4, op(e3, e4))) | (e4 = op(e3, op(e3, e4)))) & (~ (e2 = op(e4, op(e2, e4))) | (e4 = op(e2, op(e2, e4)))) & (~ (e1 = op(e4, op(e1, e4))) | (e4 = op(e1, op(e1, e4)))) & (~ (e0 = op(e4, op(e0, e4))) | (e4 = op(e0, op(e0, e4))))) | ~ sP4), inference(nnf_transformation, [], [f11])).
fof(f11, plain, (((~ (e4 = op(e4, op(e4, e4))) | (e4 = op(e4, op(e4, e4)))) & (~ (e3 = op(e4, op(e3, e4))) | (e4 = op(e3, op(e3, e4)))) & (~ (e2 = op(e4, op(e2, e4))) | (e4 = op(e2, op(e2, e4)))) & (~ (e1 = op(e4, op(e1, e4))) | (e4 = op(e1, op(e1, e4)))) & (~ (e0 = op(e4, op(e0, e4))) | (e4 = op(e0, op(e0, e4))))) | ~ sP4), inference(usedef, [], [e11])).
fof(e11, plain, (sP4 <=> ((~ (e4 = op(e4, op(e4, e4))) | (e4 = op(e4, op(e4, e4)))) & (~ (e3 = op(e4, op(e3, e4))) | (e4 = op(e3, op(e3, e4)))) & (~ (e2 = op(e4, op(e2, e4))) | (e4 = op(e2, op(e2, e4)))) & (~ (e1 = op(e4, op(e1, e4))) | (e4 = op(e1, op(e1, e4)))) & (~ (e0 = op(e4, op(e0, e4))) | (e4 = op(e0, op(e0, e4)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f1527, plain, (sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1526, f1048])).
fof(f1048, plain, ~ sP5, inference(subsumption_resolution, [], [f808, f625])).
fof(f625, plain, (unit = op(unit, unit)), inference(definition_unfolding, [], [f178, f203, f203, f203])).
fof(f178, plain, (e0 = op(e0, e0)), inference(cnf_transformation, [], [f2])).
fof(f808, plain, (~ (unit = op(unit, unit)) | ~ sP5), inference(definition_unfolding, [], [f574, f203, f203, f203])).
fof(f574, plain, (~ (e0 = op(e0, e0)) | ~ sP5), inference(cnf_transformation, [], [f162])).
fof(f162, plain, ((~ (e4 = op(e0, e0)) & ~ (e3 = op(e0, e0)) & ~ (e2 = op(e0, e0)) & ~ (e1 = op(e0, e0)) & ~ (e0 = op(e0, e0))) | ~ sP5), inference(nnf_transformation, [], [f12])).
fof(f12, plain, ((~ (e4 = op(e0, e0)) & ~ (e3 = op(e0, e0)) & ~ (e2 = op(e0, e0)) & ~ (e1 = op(e0, e0)) & ~ (e0 = op(e0, e0))) | ~ sP5), inference(usedef, [], [e12])).
fof(e12, plain, (sP5 <=> (~ (e4 = op(e0, e0)) & ~ (e3 = op(e0, e0)) & ~ (e2 = op(e0, e0)) & ~ (e1 = op(e0, e0)) & ~ (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f1526, plain, (sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1525, f1043])).
fof(f1043, plain, ~ sP6, inference(subsumption_resolution, [], [f802, f624])).
fof(f624, plain, (e1 = op(unit, e1)), inference(definition_unfolding, [], [f179, f203])).
fof(f179, plain, (e1 = op(e0, e1)), inference(cnf_transformation, [], [f2])).
fof(f802, plain, (~ (e1 = op(unit, e1)) | ~ sP6), inference(definition_unfolding, [], [f570, f203])).
fof(f570, plain, (~ (e1 = op(e0, e1)) | ~ sP6), inference(cnf_transformation, [], [f161])).
fof(f161, plain, ((~ (e4 = op(e0, e1)) & ~ (e3 = op(e0, e1)) & ~ (e2 = op(e0, e1)) & ~ (e1 = op(e0, e1)) & ~ (e0 = op(e0, e1))) | ~ sP6), inference(nnf_transformation, [], [f13])).
fof(f13, plain, ((~ (e4 = op(e0, e1)) & ~ (e3 = op(e0, e1)) & ~ (e2 = op(e0, e1)) & ~ (e1 = op(e0, e1)) & ~ (e0 = op(e0, e1))) | ~ sP6), inference(usedef, [], [e13])).
fof(e13, plain, (sP6 <=> (~ (e4 = op(e0, e1)) & ~ (e3 = op(e0, e1)) & ~ (e2 = op(e0, e1)) & ~ (e1 = op(e0, e1)) & ~ (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f1525, plain, (sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1524, f1039])).
fof(f1039, plain, ~ sP7, inference(subsumption_resolution, [], [f796, f623])).
fof(f623, plain, (e2 = op(unit, e2)), inference(definition_unfolding, [], [f180, f203])).
fof(f180, plain, (e2 = op(e0, e2)), inference(cnf_transformation, [], [f2])).
fof(f796, plain, (~ (e2 = op(unit, e2)) | ~ sP7), inference(definition_unfolding, [], [f566, f203])).
fof(f566, plain, (~ (e2 = op(e0, e2)) | ~ sP7), inference(cnf_transformation, [], [f160])).
fof(f160, plain, ((~ (e4 = op(e0, e2)) & ~ (e3 = op(e0, e2)) & ~ (e2 = op(e0, e2)) & ~ (e1 = op(e0, e2)) & ~ (e0 = op(e0, e2))) | ~ sP7), inference(nnf_transformation, [], [f14])).
fof(f14, plain, ((~ (e4 = op(e0, e2)) & ~ (e3 = op(e0, e2)) & ~ (e2 = op(e0, e2)) & ~ (e1 = op(e0, e2)) & ~ (e0 = op(e0, e2))) | ~ sP7), inference(usedef, [], [e14])).
fof(e14, plain, (sP7 <=> (~ (e4 = op(e0, e2)) & ~ (e3 = op(e0, e2)) & ~ (e2 = op(e0, e2)) & ~ (e1 = op(e0, e2)) & ~ (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f1524, plain, (sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1523, f1036])).
fof(f1036, plain, ~ sP8, inference(subsumption_resolution, [], [f790, f622])).
fof(f790, plain, (~ (e3 = op(unit, e3)) | ~ sP8), inference(definition_unfolding, [], [f562, f203])).
fof(f562, plain, (~ (e3 = op(e0, e3)) | ~ sP8), inference(cnf_transformation, [], [f159])).
fof(f159, plain, ((~ (e4 = op(e0, e3)) & ~ (e3 = op(e0, e3)) & ~ (e2 = op(e0, e3)) & ~ (e1 = op(e0, e3)) & ~ (e0 = op(e0, e3))) | ~ sP8), inference(nnf_transformation, [], [f15])).
fof(f15, plain, ((~ (e4 = op(e0, e3)) & ~ (e3 = op(e0, e3)) & ~ (e2 = op(e0, e3)) & ~ (e1 = op(e0, e3)) & ~ (e0 = op(e0, e3))) | ~ sP8), inference(usedef, [], [e15])).
fof(e15, plain, (sP8 <=> (~ (e4 = op(e0, e3)) & ~ (e3 = op(e0, e3)) & ~ (e2 = op(e0, e3)) & ~ (e1 = op(e0, e3)) & ~ (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f1523, plain, (sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1522, f1034])).
fof(f1034, plain, ~ sP9, inference(subsumption_resolution, [], [f784, f621])).
fof(f621, plain, (e4 = op(unit, e4)), inference(definition_unfolding, [], [f182, f203])).
fof(f182, plain, (e4 = op(e0, e4)), inference(cnf_transformation, [], [f2])).
fof(f784, plain, (~ (e4 = op(unit, e4)) | ~ sP9), inference(definition_unfolding, [], [f558, f203])).
fof(f558, plain, (~ (e4 = op(e0, e4)) | ~ sP9), inference(cnf_transformation, [], [f158])).
fof(f158, plain, ((~ (e4 = op(e0, e4)) & ~ (e3 = op(e0, e4)) & ~ (e2 = op(e0, e4)) & ~ (e1 = op(e0, e4)) & ~ (e0 = op(e0, e4))) | ~ sP9), inference(nnf_transformation, [], [f16])).
fof(f16, plain, ((~ (e4 = op(e0, e4)) & ~ (e3 = op(e0, e4)) & ~ (e2 = op(e0, e4)) & ~ (e1 = op(e0, e4)) & ~ (e0 = op(e0, e4))) | ~ sP9), inference(usedef, [], [e16])).
fof(e16, plain, (sP9 <=> (~ (e4 = op(e0, e4)) & ~ (e3 = op(e0, e4)) & ~ (e2 = op(e0, e4)) & ~ (e1 = op(e0, e4)) & ~ (e0 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f1522, plain, (sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1521, f1033])).
fof(f1033, plain, ~ sP10, inference(subsumption_resolution, [], [f782, f620])).
fof(f620, plain, (e1 = op(e1, unit)), inference(definition_unfolding, [], [f183, f203])).
fof(f183, plain, (e1 = op(e1, e0)), inference(cnf_transformation, [], [f2])).
fof(f782, plain, (~ (e1 = op(e1, unit)) | ~ sP10), inference(definition_unfolding, [], [f550, f203])).
fof(f550, plain, (~ (e1 = op(e1, e0)) | ~ sP10), inference(cnf_transformation, [], [f157])).
fof(f157, plain, ((~ (e4 = op(e1, e0)) & ~ (e3 = op(e1, e0)) & ~ (e2 = op(e1, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e1, e0))) | ~ sP10), inference(nnf_transformation, [], [f17])).
fof(f17, plain, ((~ (e4 = op(e1, e0)) & ~ (e3 = op(e1, e0)) & ~ (e2 = op(e1, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e1, e0))) | ~ sP10), inference(usedef, [], [e17])).
fof(e17, plain, (sP10 <=> (~ (e4 = op(e1, e0)) & ~ (e3 = op(e1, e0)) & ~ (e2 = op(e1, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f1521, plain, (sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1520, f1029])).
fof(f1029, plain, ~ sP11, inference(subsumption_resolution, [], [f548, f184])).
fof(f184, plain, (e4 = op(e1, e1)), inference(cnf_transformation, [], [f2])).
fof(f548, plain, (~ (e4 = op(e1, e1)) | ~ sP11), inference(cnf_transformation, [], [f156])).
fof(f156, plain, ((~ (e4 = op(e1, e1)) & ~ (e3 = op(e1, e1)) & ~ (e2 = op(e1, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e1, e1))) | ~ sP11), inference(nnf_transformation, [], [f18])).
fof(f18, plain, ((~ (e4 = op(e1, e1)) & ~ (e3 = op(e1, e1)) & ~ (e2 = op(e1, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e1, e1))) | ~ sP11), inference(usedef, [], [e18])).
fof(e18, plain, (sP11 <=> (~ (e4 = op(e1, e1)) & ~ (e3 = op(e1, e1)) & ~ (e2 = op(e1, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f1520, plain, (sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1519, f1028])).
fof(f1028, plain, ~ sP12, inference(subsumption_resolution, [], [f542, f185])).
fof(f542, plain, (~ (e3 = op(e1, e2)) | ~ sP12), inference(cnf_transformation, [], [f155])).
fof(f155, plain, ((~ (e4 = op(e1, e2)) & ~ (e3 = op(e1, e2)) & ~ (e2 = op(e1, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e1, e2))) | ~ sP12), inference(nnf_transformation, [], [f19])).
fof(f19, plain, ((~ (e4 = op(e1, e2)) & ~ (e3 = op(e1, e2)) & ~ (e2 = op(e1, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e1, e2))) | ~ sP12), inference(usedef, [], [e19])).
fof(e19, plain, (sP12 <=> (~ (e4 = op(e1, e2)) & ~ (e3 = op(e1, e2)) & ~ (e2 = op(e1, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f1519, plain, (sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1518, f1026])).
fof(f1026, plain, ~ sP13, inference(subsumption_resolution, [], [f776, f619])).
fof(f776, plain, (~ (op(e1, e3) = unit) | ~ sP13), inference(definition_unfolding, [], [f534, f203])).
fof(f534, plain, (~ (e0 = op(e1, e3)) | ~ sP13), inference(cnf_transformation, [], [f154])).
fof(f154, plain, ((~ (e4 = op(e1, e3)) & ~ (e3 = op(e1, e3)) & ~ (e2 = op(e1, e3)) & ~ (e1 = op(e1, e3)) & ~ (e0 = op(e1, e3))) | ~ sP13), inference(nnf_transformation, [], [f20])).
fof(f20, plain, ((~ (e4 = op(e1, e3)) & ~ (e3 = op(e1, e3)) & ~ (e2 = op(e1, e3)) & ~ (e1 = op(e1, e3)) & ~ (e0 = op(e1, e3))) | ~ sP13), inference(usedef, [], [e20])).
fof(e20, plain, (sP13 <=> (~ (e4 = op(e1, e3)) & ~ (e3 = op(e1, e3)) & ~ (e2 = op(e1, e3)) & ~ (e1 = op(e1, e3)) & ~ (e0 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f1518, plain, (sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1517, f1021])).
fof(f1021, plain, ~ sP14, inference(subsumption_resolution, [], [f531, f187])).
fof(f187, plain, (e2 = op(e1, e4)), inference(cnf_transformation, [], [f2])).
fof(f531, plain, (~ (e2 = op(e1, e4)) | ~ sP14), inference(cnf_transformation, [], [f153])).
fof(f153, plain, ((~ (e4 = op(e1, e4)) & ~ (e3 = op(e1, e4)) & ~ (e2 = op(e1, e4)) & ~ (e1 = op(e1, e4)) & ~ (e0 = op(e1, e4))) | ~ sP14), inference(nnf_transformation, [], [f21])).
fof(f21, plain, ((~ (e4 = op(e1, e4)) & ~ (e3 = op(e1, e4)) & ~ (e2 = op(e1, e4)) & ~ (e1 = op(e1, e4)) & ~ (e0 = op(e1, e4))) | ~ sP14), inference(usedef, [], [e21])).
fof(e21, plain, (sP14 <=> (~ (e4 = op(e1, e4)) & ~ (e3 = op(e1, e4)) & ~ (e2 = op(e1, e4)) & ~ (e1 = op(e1, e4)) & ~ (e0 = op(e1, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f1517, plain, (sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1516, f1018])).
fof(f1018, plain, ~ sP15, inference(subsumption_resolution, [], [f772, f618])).
fof(f618, plain, (e2 = op(e2, unit)), inference(definition_unfolding, [], [f188, f203])).
fof(f188, plain, (e2 = op(e2, e0)), inference(cnf_transformation, [], [f2])).
fof(f772, plain, (~ (e2 = op(e2, unit)) | ~ sP15), inference(definition_unfolding, [], [f526, f203])).
fof(f526, plain, (~ (e2 = op(e2, e0)) | ~ sP15), inference(cnf_transformation, [], [f152])).
fof(f152, plain, ((~ (e4 = op(e2, e0)) & ~ (e3 = op(e2, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e2, e0)) & ~ (e0 = op(e2, e0))) | ~ sP15), inference(nnf_transformation, [], [f22])).
fof(f22, plain, ((~ (e4 = op(e2, e0)) & ~ (e3 = op(e2, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e2, e0)) & ~ (e0 = op(e2, e0))) | ~ sP15), inference(usedef, [], [e22])).
fof(e22, plain, (sP15 <=> (~ (e4 = op(e2, e0)) & ~ (e3 = op(e2, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e2, e0)) & ~ (e0 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f1516, plain, (sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1515, f1015])).
fof(f1015, plain, ~ sP16, inference(subsumption_resolution, [], [f769, f617])).
fof(f769, plain, (~ (op(e2, e1) = unit) | ~ sP16), inference(definition_unfolding, [], [f519, f203])).
fof(f519, plain, (~ (e0 = op(e2, e1)) | ~ sP16), inference(cnf_transformation, [], [f151])).
fof(f151, plain, ((~ (e4 = op(e2, e1)) & ~ (e3 = op(e2, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e2, e1)) & ~ (e0 = op(e2, e1))) | ~ sP16), inference(nnf_transformation, [], [f23])).
fof(f23, plain, ((~ (e4 = op(e2, e1)) & ~ (e3 = op(e2, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e2, e1)) & ~ (e0 = op(e2, e1))) | ~ sP16), inference(usedef, [], [e23])).
fof(e23, plain, (sP16 <=> (~ (e4 = op(e2, e1)) & ~ (e3 = op(e2, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e2, e1)) & ~ (e0 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f1515, plain, (sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1514, f1010])).
fof(f1010, plain, ~ sP17, inference(subsumption_resolution, [], [f518, f190])).
fof(f190, plain, (e4 = op(e2, e2)), inference(cnf_transformation, [], [f2])).
fof(f518, plain, (~ (e4 = op(e2, e2)) | ~ sP17), inference(cnf_transformation, [], [f150])).
fof(f150, plain, ((~ (e4 = op(e2, e2)) & ~ (e3 = op(e2, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e2, e2)) & ~ (e0 = op(e2, e2))) | ~ sP17), inference(nnf_transformation, [], [f24])).
fof(f24, plain, ((~ (e4 = op(e2, e2)) & ~ (e3 = op(e2, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e2, e2)) & ~ (e0 = op(e2, e2))) | ~ sP17), inference(usedef, [], [e24])).
fof(e24, plain, (sP17 <=> (~ (e4 = op(e2, e2)) & ~ (e3 = op(e2, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e2, e2)) & ~ (e0 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f1514, plain, (sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1513, f1009])).
fof(f1009, plain, ~ sP18, inference(subsumption_resolution, [], [f510, f191])).
fof(f510, plain, (~ (e1 = op(e2, e3)) | ~ sP18), inference(cnf_transformation, [], [f149])).
fof(f149, plain, ((~ (e4 = op(e2, e3)) & ~ (e3 = op(e2, e3)) & ~ (e2 = op(e2, e3)) & ~ (e1 = op(e2, e3)) & ~ (e0 = op(e2, e3))) | ~ sP18), inference(nnf_transformation, [], [f25])).
fof(f25, plain, ((~ (e4 = op(e2, e3)) & ~ (e3 = op(e2, e3)) & ~ (e2 = op(e2, e3)) & ~ (e1 = op(e2, e3)) & ~ (e0 = op(e2, e3))) | ~ sP18), inference(usedef, [], [e25])).
fof(e25, plain, (sP18 <=> (~ (e4 = op(e2, e3)) & ~ (e3 = op(e2, e3)) & ~ (e2 = op(e2, e3)) & ~ (e1 = op(e2, e3)) & ~ (e0 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f1513, plain, (sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1512, f1005])).
fof(f1005, plain, ~ sP19, inference(subsumption_resolution, [], [f507, f192])).
fof(f192, plain, (e3 = op(e2, e4)), inference(cnf_transformation, [], [f2])).
fof(f507, plain, (~ (e3 = op(e2, e4)) | ~ sP19), inference(cnf_transformation, [], [f148])).
fof(f148, plain, ((~ (e4 = op(e2, e4)) & ~ (e3 = op(e2, e4)) & ~ (e2 = op(e2, e4)) & ~ (e1 = op(e2, e4)) & ~ (e0 = op(e2, e4))) | ~ sP19), inference(nnf_transformation, [], [f26])).
fof(f26, plain, ((~ (e4 = op(e2, e4)) & ~ (e3 = op(e2, e4)) & ~ (e2 = op(e2, e4)) & ~ (e1 = op(e2, e4)) & ~ (e0 = op(e2, e4))) | ~ sP19), inference(usedef, [], [e26])).
fof(e26, plain, (sP19 <=> (~ (e4 = op(e2, e4)) & ~ (e3 = op(e2, e4)) & ~ (e2 = op(e2, e4)) & ~ (e1 = op(e2, e4)) & ~ (e0 = op(e2, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP19])])).
fof(f1512, plain, (sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1511, f1003])).
fof(f1003, plain, ~ sP20, inference(subsumption_resolution, [], [f762, f616])).
fof(f762, plain, (~ (e3 = op(e3, unit)) | ~ sP20), inference(definition_unfolding, [], [f502, f203])).
fof(f502, plain, (~ (e3 = op(e3, e0)) | ~ sP20), inference(cnf_transformation, [], [f147])).
fof(f147, plain, ((~ (e4 = op(e3, e0)) & ~ (e3 = op(e3, e0)) & ~ (e2 = op(e3, e0)) & ~ (e1 = op(e3, e0)) & ~ (e0 = op(e3, e0))) | ~ sP20), inference(nnf_transformation, [], [f27])).
fof(f27, plain, ((~ (e4 = op(e3, e0)) & ~ (e3 = op(e3, e0)) & ~ (e2 = op(e3, e0)) & ~ (e1 = op(e3, e0)) & ~ (e0 = op(e3, e0))) | ~ sP20), inference(usedef, [], [e27])).
fof(e27, plain, (sP20 <=> (~ (e4 = op(e3, e0)) & ~ (e3 = op(e3, e0)) & ~ (e2 = op(e3, e0)) & ~ (e1 = op(e3, e0)) & ~ (e0 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP20])])).
fof(f1511, plain, (sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1510, f1001])).
fof(f1001, plain, ~ sP21, inference(subsumption_resolution, [], [f496, f194])).
fof(f496, plain, (~ (e2 = op(e3, e1)) | ~ sP21), inference(cnf_transformation, [], [f146])).
fof(f146, plain, ((~ (e4 = op(e3, e1)) & ~ (e3 = op(e3, e1)) & ~ (e2 = op(e3, e1)) & ~ (e1 = op(e3, e1)) & ~ (e0 = op(e3, e1))) | ~ sP21), inference(nnf_transformation, [], [f28])).
fof(f28, plain, ((~ (e4 = op(e3, e1)) & ~ (e3 = op(e3, e1)) & ~ (e2 = op(e3, e1)) & ~ (e1 = op(e3, e1)) & ~ (e0 = op(e3, e1))) | ~ sP21), inference(usedef, [], [e28])).
fof(e28, plain, (sP21 <=> (~ (e4 = op(e3, e1)) & ~ (e3 = op(e3, e1)) & ~ (e2 = op(e3, e1)) & ~ (e1 = op(e3, e1)) & ~ (e0 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP21])])).
fof(f1510, plain, (sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1509, f998])).
fof(f998, plain, ~ sP22, inference(subsumption_resolution, [], [f759, f615])).
fof(f759, plain, (~ (op(e3, e2) = unit) | ~ sP22), inference(definition_unfolding, [], [f489, f203])).
fof(f489, plain, (~ (e0 = op(e3, e2)) | ~ sP22), inference(cnf_transformation, [], [f145])).
fof(f145, plain, ((~ (e4 = op(e3, e2)) & ~ (e3 = op(e3, e2)) & ~ (e2 = op(e3, e2)) & ~ (e1 = op(e3, e2)) & ~ (e0 = op(e3, e2))) | ~ sP22), inference(nnf_transformation, [], [f29])).
fof(f29, plain, ((~ (e4 = op(e3, e2)) & ~ (e3 = op(e3, e2)) & ~ (e2 = op(e3, e2)) & ~ (e1 = op(e3, e2)) & ~ (e0 = op(e3, e2))) | ~ sP22), inference(usedef, [], [e29])).
fof(e29, plain, (sP22 <=> (~ (e4 = op(e3, e2)) & ~ (e3 = op(e3, e2)) & ~ (e2 = op(e3, e2)) & ~ (e1 = op(e3, e2)) & ~ (e0 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP22])])).
fof(f1509, plain, (sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1508, f993])).
fof(f993, plain, ~ sP23, inference(subsumption_resolution, [], [f488, f196])).
fof(f488, plain, (~ (e4 = op(e3, e3)) | ~ sP23), inference(cnf_transformation, [], [f144])).
fof(f144, plain, ((~ (e4 = op(e3, e3)) & ~ (e3 = op(e3, e3)) & ~ (e2 = op(e3, e3)) & ~ (e1 = op(e3, e3)) & ~ (e0 = op(e3, e3))) | ~ sP23), inference(nnf_transformation, [], [f30])).
fof(f30, plain, ((~ (e4 = op(e3, e3)) & ~ (e3 = op(e3, e3)) & ~ (e2 = op(e3, e3)) & ~ (e1 = op(e3, e3)) & ~ (e0 = op(e3, e3))) | ~ sP23), inference(usedef, [], [e30])).
fof(e30, plain, (sP23 <=> (~ (e4 = op(e3, e3)) & ~ (e3 = op(e3, e3)) & ~ (e2 = op(e3, e3)) & ~ (e1 = op(e3, e3)) & ~ (e0 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP23])])).
fof(f1508, plain, (sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1507, f992])).
fof(f992, plain, ~ sP24, inference(subsumption_resolution, [], [f480, f197])).
fof(f480, plain, (~ (e1 = op(e3, e4)) | ~ sP24), inference(cnf_transformation, [], [f143])).
fof(f143, plain, ((~ (e4 = op(e3, e4)) & ~ (e3 = op(e3, e4)) & ~ (e2 = op(e3, e4)) & ~ (e1 = op(e3, e4)) & ~ (e0 = op(e3, e4))) | ~ sP24), inference(nnf_transformation, [], [f31])).
fof(f31, plain, ((~ (e4 = op(e3, e4)) & ~ (e3 = op(e3, e4)) & ~ (e2 = op(e3, e4)) & ~ (e1 = op(e3, e4)) & ~ (e0 = op(e3, e4))) | ~ sP24), inference(usedef, [], [e31])).
fof(e31, plain, (sP24 <=> (~ (e4 = op(e3, e4)) & ~ (e3 = op(e3, e4)) & ~ (e2 = op(e3, e4)) & ~ (e1 = op(e3, e4)) & ~ (e0 = op(e3, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP24])])).
fof(f1507, plain, (sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1506, f988])).
fof(f988, plain, ~ sP25, inference(subsumption_resolution, [], [f752, f614])).
fof(f614, plain, (e4 = op(e4, unit)), inference(definition_unfolding, [], [f198, f203])).
fof(f198, plain, (e4 = op(e4, e0)), inference(cnf_transformation, [], [f2])).
fof(f752, plain, (~ (e4 = op(e4, unit)) | ~ sP25), inference(definition_unfolding, [], [f478, f203])).
fof(f478, plain, (~ (e4 = op(e4, e0)) | ~ sP25), inference(cnf_transformation, [], [f142])).
fof(f142, plain, ((~ (e4 = op(e4, e0)) & ~ (e3 = op(e4, e0)) & ~ (e2 = op(e4, e0)) & ~ (e1 = op(e4, e0)) & ~ (e0 = op(e4, e0))) | ~ sP25), inference(nnf_transformation, [], [f32])).
fof(f32, plain, ((~ (e4 = op(e4, e0)) & ~ (e3 = op(e4, e0)) & ~ (e2 = op(e4, e0)) & ~ (e1 = op(e4, e0)) & ~ (e0 = op(e4, e0))) | ~ sP25), inference(usedef, [], [e32])).
fof(e32, plain, (sP25 <=> (~ (e4 = op(e4, e0)) & ~ (e3 = op(e4, e0)) & ~ (e2 = op(e4, e0)) & ~ (e1 = op(e4, e0)) & ~ (e0 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP25])])).
fof(f1506, plain, (sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1505, f987])).
fof(f987, plain, ~ sP26, inference(subsumption_resolution, [], [f472, f199])).
fof(f472, plain, (~ (e3 = op(e4, e1)) | ~ sP26), inference(cnf_transformation, [], [f141])).
fof(f141, plain, ((~ (e4 = op(e4, e1)) & ~ (e3 = op(e4, e1)) & ~ (e2 = op(e4, e1)) & ~ (e1 = op(e4, e1)) & ~ (e0 = op(e4, e1))) | ~ sP26), inference(nnf_transformation, [], [f33])).
fof(f33, plain, ((~ (e4 = op(e4, e1)) & ~ (e3 = op(e4, e1)) & ~ (e2 = op(e4, e1)) & ~ (e1 = op(e4, e1)) & ~ (e0 = op(e4, e1))) | ~ sP26), inference(usedef, [], [e33])).
fof(e33, plain, (sP26 <=> (~ (e4 = op(e4, e1)) & ~ (e3 = op(e4, e1)) & ~ (e2 = op(e4, e1)) & ~ (e1 = op(e4, e1)) & ~ (e0 = op(e4, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP26])])).
fof(f1505, plain, (sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1504, f985])).
fof(f985, plain, ~ sP27, inference(subsumption_resolution, [], [f465, f200])).
fof(f200, plain, (e1 = op(e4, e2)), inference(cnf_transformation, [], [f2])).
fof(f465, plain, (~ (e1 = op(e4, e2)) | ~ sP27), inference(cnf_transformation, [], [f140])).
fof(f140, plain, ((~ (e4 = op(e4, e2)) & ~ (e3 = op(e4, e2)) & ~ (e2 = op(e4, e2)) & ~ (e1 = op(e4, e2)) & ~ (e0 = op(e4, e2))) | ~ sP27), inference(nnf_transformation, [], [f34])).
fof(f34, plain, ((~ (e4 = op(e4, e2)) & ~ (e3 = op(e4, e2)) & ~ (e2 = op(e4, e2)) & ~ (e1 = op(e4, e2)) & ~ (e0 = op(e4, e2))) | ~ sP27), inference(usedef, [], [e34])).
fof(e34, plain, (sP27 <=> (~ (e4 = op(e4, e2)) & ~ (e3 = op(e4, e2)) & ~ (e2 = op(e4, e2)) & ~ (e1 = op(e4, e2)) & ~ (e0 = op(e4, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP27])])).
fof(f1504, plain, (sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1503, f981])).
fof(f981, plain, ~ sP28, inference(subsumption_resolution, [], [f461, f201])).
fof(f201, plain, (e2 = op(e4, e3)), inference(cnf_transformation, [], [f2])).
fof(f461, plain, (~ (e2 = op(e4, e3)) | ~ sP28), inference(cnf_transformation, [], [f139])).
fof(f139, plain, ((~ (e4 = op(e4, e3)) & ~ (e3 = op(e4, e3)) & ~ (e2 = op(e4, e3)) & ~ (e1 = op(e4, e3)) & ~ (e0 = op(e4, e3))) | ~ sP28), inference(nnf_transformation, [], [f35])).
fof(f35, plain, ((~ (e4 = op(e4, e3)) & ~ (e3 = op(e4, e3)) & ~ (e2 = op(e4, e3)) & ~ (e1 = op(e4, e3)) & ~ (e0 = op(e4, e3))) | ~ sP28), inference(usedef, [], [e35])).
fof(e35, plain, (sP28 <=> (~ (e4 = op(e4, e3)) & ~ (e3 = op(e4, e3)) & ~ (e2 = op(e4, e3)) & ~ (e1 = op(e4, e3)) & ~ (e0 = op(e4, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP28])])).
fof(f1503, plain, (sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1502, f978])).
fof(f978, plain, ~ sP29, inference(subsumption_resolution, [], [f748, f613])).
fof(f613, plain, (op(e4, e4) = unit), inference(definition_unfolding, [], [f202, f203])).
fof(f202, plain, (e0 = op(e4, e4)), inference(cnf_transformation, [], [f2])).
fof(f748, plain, (~ (op(e4, e4) = unit) | ~ sP29), inference(definition_unfolding, [], [f454, f203])).
fof(f454, plain, (~ (e0 = op(e4, e4)) | ~ sP29), inference(cnf_transformation, [], [f138])).
fof(f138, plain, ((~ (e4 = op(e4, e4)) & ~ (e3 = op(e4, e4)) & ~ (e2 = op(e4, e4)) & ~ (e1 = op(e4, e4)) & ~ (e0 = op(e4, e4))) | ~ sP29), inference(nnf_transformation, [], [f36])).
fof(f36, plain, ((~ (e4 = op(e4, e4)) & ~ (e3 = op(e4, e4)) & ~ (e2 = op(e4, e4)) & ~ (e1 = op(e4, e4)) & ~ (e0 = op(e4, e4))) | ~ sP29), inference(usedef, [], [e36])).
fof(e36, plain, (sP29 <=> (~ (e4 = op(e4, e4)) & ~ (e3 = op(e4, e4)) & ~ (e2 = op(e4, e4)) & ~ (e1 = op(e4, e4)) & ~ (e0 = op(e4, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP29])])).
fof(f1502, plain, (sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1501, f625])).
fof(f1501, plain, (~ (unit = op(unit, unit)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1500, f624])).
fof(f1500, plain, (~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1499, f620])).
fof(f1499, plain, (~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1498, f623])).
fof(f1498, plain, (~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1497, f618])).
fof(f1497, plain, (~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1496, f622])).
fof(f1496, plain, (~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1495, f616])).
fof(f1495, plain, (~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1494, f614])).
fof(f1494, plain, (~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1493, f828])).
fof(f828, plain, ~ sP30, inference(trivial_inequality_removal, [], [f747])).
fof(f747, plain, (~ (unit = unit) | ~ sP30), inference(definition_unfolding, [], [f449, f203])).
fof(f449, plain, (~ (e0 = unit) | ~ sP30), inference(cnf_transformation, [], [f137])).
fof(f137, plain, ((~ (e4 = unit) & ~ (e3 = unit) & ~ (e2 = unit) & ~ (e1 = unit) & ~ (e0 = unit)) | ~ sP30), inference(nnf_transformation, [], [f37])).
fof(f37, plain, ((~ (e4 = unit) & ~ (e3 = unit) & ~ (e2 = unit) & ~ (e1 = unit) & ~ (e0 = unit)) | ~ sP30), inference(usedef, [], [e37])).
fof(e37, plain, (sP30 <=> (~ (e4 = unit) & ~ (e3 = unit) & ~ (e2 = unit) & ~ (e1 = unit) & ~ (e0 = unit))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP30])])).
fof(f1493, plain, (sP30 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1492, f973])).
fof(f973, plain, ~ sP31, inference(subsumption_resolution, [], [f746, f625])).
fof(f746, plain, (~ (unit = op(unit, unit)) | ~ sP31), inference(definition_unfolding, [], [f444, f203, f203, f203])).
fof(f444, plain, (~ (e0 = op(e0, e0)) | ~ sP31), inference(cnf_transformation, [], [f136])).
fof(f136, plain, ((~ (e0 = op(e0, e4)) & ~ (e0 = op(e0, e3)) & ~ (e0 = op(e0, e2)) & ~ (e0 = op(e0, e1)) & ~ (e0 = op(e0, e0))) | ~ sP31), inference(nnf_transformation, [], [f38])).
fof(f38, plain, ((~ (e0 = op(e0, e4)) & ~ (e0 = op(e0, e3)) & ~ (e0 = op(e0, e2)) & ~ (e0 = op(e0, e1)) & ~ (e0 = op(e0, e0))) | ~ sP31), inference(usedef, [], [e38])).
fof(e38, plain, (sP31 <=> (~ (e0 = op(e0, e4)) & ~ (e0 = op(e0, e3)) & ~ (e0 = op(e0, e2)) & ~ (e0 = op(e0, e1)) & ~ (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP31])])).
fof(f1492, plain, (sP31 | sP30 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1491, f968])).
fof(f968, plain, ~ sP32, inference(subsumption_resolution, [], [f741, f625])).
fof(f741, plain, (~ (unit = op(unit, unit)) | ~ sP32), inference(definition_unfolding, [], [f439, f203, f203, f203])).
fof(f439, plain, (~ (e0 = op(e0, e0)) | ~ sP32), inference(cnf_transformation, [], [f135])).
fof(f135, plain, ((~ (e0 = op(e4, e0)) & ~ (e0 = op(e3, e0)) & ~ (e0 = op(e2, e0)) & ~ (e0 = op(e1, e0)) & ~ (e0 = op(e0, e0))) | ~ sP32), inference(nnf_transformation, [], [f39])).
fof(f39, plain, ((~ (e0 = op(e4, e0)) & ~ (e0 = op(e3, e0)) & ~ (e0 = op(e2, e0)) & ~ (e0 = op(e1, e0)) & ~ (e0 = op(e0, e0))) | ~ sP32), inference(usedef, [], [e39])).
fof(e39, plain, (sP32 <=> (~ (e0 = op(e4, e0)) & ~ (e0 = op(e3, e0)) & ~ (e0 = op(e2, e0)) & ~ (e0 = op(e1, e0)) & ~ (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP32])])).
fof(f1491, plain, (sP32 | sP31 | sP30 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1490, f963])).
fof(f963, plain, ~ sP33, inference(subsumption_resolution, [], [f735, f624])).
fof(f735, plain, (~ (e1 = op(unit, e1)) | ~ sP33), inference(definition_unfolding, [], [f435, f203])).
fof(f435, plain, (~ (e1 = op(e0, e1)) | ~ sP33), inference(cnf_transformation, [], [f134])).
fof(f134, plain, ((~ (e1 = op(e0, e4)) & ~ (e1 = op(e0, e3)) & ~ (e1 = op(e0, e2)) & ~ (e1 = op(e0, e1)) & ~ (e1 = op(e0, e0))) | ~ sP33), inference(nnf_transformation, [], [f40])).
fof(f40, plain, ((~ (e1 = op(e0, e4)) & ~ (e1 = op(e0, e3)) & ~ (e1 = op(e0, e2)) & ~ (e1 = op(e0, e1)) & ~ (e1 = op(e0, e0))) | ~ sP33), inference(usedef, [], [e40])).
fof(e40, plain, (sP33 <=> (~ (e1 = op(e0, e4)) & ~ (e1 = op(e0, e3)) & ~ (e1 = op(e0, e2)) & ~ (e1 = op(e0, e1)) & ~ (e1 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP33])])).
fof(f1490, plain, (sP33 | sP32 | sP31 | sP30 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1489, f959])).
fof(f959, plain, ~ sP34, inference(subsumption_resolution, [], [f730, f620])).
fof(f730, plain, (~ (e1 = op(e1, unit)) | ~ sP34), inference(definition_unfolding, [], [f430, f203])).
fof(f430, plain, (~ (e1 = op(e1, e0)) | ~ sP34), inference(cnf_transformation, [], [f133])).
fof(f133, plain, ((~ (e1 = op(e4, e0)) & ~ (e1 = op(e3, e0)) & ~ (e1 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e1 = op(e0, e0))) | ~ sP34), inference(nnf_transformation, [], [f41])).
fof(f41, plain, ((~ (e1 = op(e4, e0)) & ~ (e1 = op(e3, e0)) & ~ (e1 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e1 = op(e0, e0))) | ~ sP34), inference(usedef, [], [e41])).
fof(e41, plain, (sP34 <=> (~ (e1 = op(e4, e0)) & ~ (e1 = op(e3, e0)) & ~ (e1 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e1 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP34])])).
fof(f1489, plain, (sP34 | sP33 | sP32 | sP31 | sP30 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1488, f955])).
fof(f955, plain, ~ sP35, inference(subsumption_resolution, [], [f724, f623])).
fof(f724, plain, (~ (e2 = op(unit, e2)) | ~ sP35), inference(definition_unfolding, [], [f426, f203])).
fof(f426, plain, (~ (e2 = op(e0, e2)) | ~ sP35), inference(cnf_transformation, [], [f132])).
fof(f132, plain, ((~ (e2 = op(e0, e4)) & ~ (e2 = op(e0, e3)) & ~ (e2 = op(e0, e2)) & ~ (e2 = op(e0, e1)) & ~ (e2 = op(e0, e0))) | ~ sP35), inference(nnf_transformation, [], [f42])).
fof(f42, plain, ((~ (e2 = op(e0, e4)) & ~ (e2 = op(e0, e3)) & ~ (e2 = op(e0, e2)) & ~ (e2 = op(e0, e1)) & ~ (e2 = op(e0, e0))) | ~ sP35), inference(usedef, [], [e42])).
fof(e42, plain, (sP35 <=> (~ (e2 = op(e0, e4)) & ~ (e2 = op(e0, e3)) & ~ (e2 = op(e0, e2)) & ~ (e2 = op(e0, e1)) & ~ (e2 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP35])])).
fof(f1488, plain, (sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1487, f952])).
fof(f952, plain, ~ sP36, inference(subsumption_resolution, [], [f719, f618])).
fof(f719, plain, (~ (e2 = op(e2, unit)) | ~ sP36), inference(definition_unfolding, [], [f421, f203])).
fof(f421, plain, (~ (e2 = op(e2, e0)) | ~ sP36), inference(cnf_transformation, [], [f131])).
fof(f131, plain, ((~ (e2 = op(e4, e0)) & ~ (e2 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e2 = op(e1, e0)) & ~ (e2 = op(e0, e0))) | ~ sP36), inference(nnf_transformation, [], [f43])).
fof(f43, plain, ((~ (e2 = op(e4, e0)) & ~ (e2 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e2 = op(e1, e0)) & ~ (e2 = op(e0, e0))) | ~ sP36), inference(usedef, [], [e43])).
fof(e43, plain, (sP36 <=> (~ (e2 = op(e4, e0)) & ~ (e2 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e2 = op(e1, e0)) & ~ (e2 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP36])])).
fof(f1487, plain, (sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1486, f949])).
fof(f949, plain, ~ sP37, inference(subsumption_resolution, [], [f713, f622])).
fof(f713, plain, (~ (e3 = op(unit, e3)) | ~ sP37), inference(definition_unfolding, [], [f417, f203])).
fof(f417, plain, (~ (e3 = op(e0, e3)) | ~ sP37), inference(cnf_transformation, [], [f130])).
fof(f130, plain, ((~ (e3 = op(e0, e4)) & ~ (e3 = op(e0, e3)) & ~ (e3 = op(e0, e2)) & ~ (e3 = op(e0, e1)) & ~ (e3 = op(e0, e0))) | ~ sP37), inference(nnf_transformation, [], [f44])).
fof(f44, plain, ((~ (e3 = op(e0, e4)) & ~ (e3 = op(e0, e3)) & ~ (e3 = op(e0, e2)) & ~ (e3 = op(e0, e1)) & ~ (e3 = op(e0, e0))) | ~ sP37), inference(usedef, [], [e44])).
fof(e44, plain, (sP37 <=> (~ (e3 = op(e0, e4)) & ~ (e3 = op(e0, e3)) & ~ (e3 = op(e0, e2)) & ~ (e3 = op(e0, e1)) & ~ (e3 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP37])])).
fof(f1486, plain, (sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1485, f947])).
fof(f947, plain, ~ sP38, inference(subsumption_resolution, [], [f708, f616])).
fof(f708, plain, (~ (e3 = op(e3, unit)) | ~ sP38), inference(definition_unfolding, [], [f412, f203])).
fof(f412, plain, (~ (e3 = op(e3, e0)) | ~ sP38), inference(cnf_transformation, [], [f129])).
fof(f129, plain, ((~ (e3 = op(e4, e0)) & ~ (e3 = op(e3, e0)) & ~ (e3 = op(e2, e0)) & ~ (e3 = op(e1, e0)) & ~ (e3 = op(e0, e0))) | ~ sP38), inference(nnf_transformation, [], [f45])).
fof(f45, plain, ((~ (e3 = op(e4, e0)) & ~ (e3 = op(e3, e0)) & ~ (e3 = op(e2, e0)) & ~ (e3 = op(e1, e0)) & ~ (e3 = op(e0, e0))) | ~ sP38), inference(usedef, [], [e45])).
fof(e45, plain, (sP38 <=> (~ (e3 = op(e4, e0)) & ~ (e3 = op(e3, e0)) & ~ (e3 = op(e2, e0)) & ~ (e3 = op(e1, e0)) & ~ (e3 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP38])])).
fof(f1485, plain, (sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1484, f945])).
fof(f945, plain, ~ sP39, inference(subsumption_resolution, [], [f702, f621])).
fof(f702, plain, (~ (e4 = op(unit, e4)) | ~ sP39), inference(definition_unfolding, [], [f408, f203])).
fof(f408, plain, (~ (e4 = op(e0, e4)) | ~ sP39), inference(cnf_transformation, [], [f128])).
fof(f128, plain, ((~ (e4 = op(e0, e4)) & ~ (e4 = op(e0, e3)) & ~ (e4 = op(e0, e2)) & ~ (e4 = op(e0, e1)) & ~ (e4 = op(e0, e0))) | ~ sP39), inference(nnf_transformation, [], [f46])).
fof(f46, plain, ((~ (e4 = op(e0, e4)) & ~ (e4 = op(e0, e3)) & ~ (e4 = op(e0, e2)) & ~ (e4 = op(e0, e1)) & ~ (e4 = op(e0, e0))) | ~ sP39), inference(usedef, [], [e46])).
fof(e46, plain, (sP39 <=> (~ (e4 = op(e0, e4)) & ~ (e4 = op(e0, e3)) & ~ (e4 = op(e0, e2)) & ~ (e4 = op(e0, e1)) & ~ (e4 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP39])])).
fof(f1484, plain, (sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1483, f944])).
fof(f944, plain, ~ sP40, inference(subsumption_resolution, [], [f697, f614])).
fof(f697, plain, (~ (e4 = op(e4, unit)) | ~ sP40), inference(definition_unfolding, [], [f403, f203])).
fof(f403, plain, (~ (e4 = op(e4, e0)) | ~ sP40), inference(cnf_transformation, [], [f127])).
fof(f127, plain, ((~ (e4 = op(e4, e0)) & ~ (e4 = op(e3, e0)) & ~ (e4 = op(e2, e0)) & ~ (e4 = op(e1, e0)) & ~ (e4 = op(e0, e0))) | ~ sP40), inference(nnf_transformation, [], [f47])).
fof(f47, plain, ((~ (e4 = op(e4, e0)) & ~ (e4 = op(e3, e0)) & ~ (e4 = op(e2, e0)) & ~ (e4 = op(e1, e0)) & ~ (e4 = op(e0, e0))) | ~ sP40), inference(usedef, [], [e47])).
fof(e47, plain, (sP40 <=> (~ (e4 = op(e4, e0)) & ~ (e4 = op(e3, e0)) & ~ (e4 = op(e2, e0)) & ~ (e4 = op(e1, e0)) & ~ (e4 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP40])])).
fof(f1483, plain, (sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1482, f943])).
fof(f943, plain, ~ sP41, inference(subsumption_resolution, [], [f693, f619])).
fof(f693, plain, (~ (op(e1, e3) = unit) | ~ sP41), inference(definition_unfolding, [], [f397, f203])).
fof(f397, plain, (~ (e0 = op(e1, e3)) | ~ sP41), inference(cnf_transformation, [], [f126])).
fof(f126, plain, ((~ (e0 = op(e1, e4)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e1, e0))) | ~ sP41), inference(nnf_transformation, [], [f48])).
fof(f48, plain, ((~ (e0 = op(e1, e4)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e1, e0))) | ~ sP41), inference(usedef, [], [e48])).
fof(e48, plain, (sP41 <=> (~ (e0 = op(e1, e4)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP41])])).
fof(f1482, plain, (sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1481, f941])).
fof(f941, plain, ~ sP42, inference(subsumption_resolution, [], [f689, f617])).
fof(f689, plain, (~ (op(e2, e1) = unit) | ~ sP42), inference(definition_unfolding, [], [f391, f203])).
fof(f391, plain, (~ (e0 = op(e2, e1)) | ~ sP42), inference(cnf_transformation, [], [f125])).
fof(f125, plain, ((~ (e0 = op(e4, e1)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e0, e1))) | ~ sP42), inference(nnf_transformation, [], [f49])).
fof(f49, plain, ((~ (e0 = op(e4, e1)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e0, e1))) | ~ sP42), inference(usedef, [], [e49])).
fof(e49, plain, (sP42 <=> (~ (e0 = op(e4, e1)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP42])])).
fof(f1481, plain, (sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1480, f938])).
fof(f938, plain, ~ sP43, inference(subsumption_resolution, [], [f686, f620])).
fof(f686, plain, (~ (e1 = op(e1, unit)) | ~ sP43), inference(definition_unfolding, [], [f384, f203])).
fof(f384, plain, (~ (e1 = op(e1, e0)) | ~ sP43), inference(cnf_transformation, [], [f124])).
fof(f124, plain, ((~ (e1 = op(e1, e4)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e1, e0))) | ~ sP43), inference(nnf_transformation, [], [f50])).
fof(f50, plain, ((~ (e1 = op(e1, e4)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e1, e0))) | ~ sP43), inference(usedef, [], [e50])).
fof(e50, plain, (sP43 <=> (~ (e1 = op(e1, e4)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP43])])).
fof(f1480, plain, (sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1479, f933])).
fof(f933, plain, ~ sP44, inference(subsumption_resolution, [], [f685, f624])).
fof(f685, plain, (~ (e1 = op(unit, e1)) | ~ sP44), inference(definition_unfolding, [], [f379, f203])).
fof(f379, plain, (~ (e1 = op(e0, e1)) | ~ sP44), inference(cnf_transformation, [], [f123])).
fof(f123, plain, ((~ (e1 = op(e4, e1)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e0, e1))) | ~ sP44), inference(nnf_transformation, [], [f51])).
fof(f51, plain, ((~ (e1 = op(e4, e1)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e0, e1))) | ~ sP44), inference(usedef, [], [e51])).
fof(e51, plain, (sP44 <=> (~ (e1 = op(e4, e1)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP44])])).
fof(f1479, plain, (sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1478, f928])).
fof(f928, plain, ~ sP45, inference(subsumption_resolution, [], [f378, f187])).
fof(f378, plain, (~ (e2 = op(e1, e4)) | ~ sP45), inference(cnf_transformation, [], [f122])).
fof(f122, plain, ((~ (e2 = op(e1, e4)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e1, e0))) | ~ sP45), inference(nnf_transformation, [], [f52])).
fof(f52, plain, ((~ (e2 = op(e1, e4)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e1, e0))) | ~ sP45), inference(usedef, [], [e52])).
fof(e52, plain, (sP45 <=> (~ (e2 = op(e1, e4)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP45])])).
fof(f1478, plain, (sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1477, f927])).
fof(f927, plain, ~ sP46, inference(subsumption_resolution, [], [f372, f194])).
fof(f372, plain, (~ (e2 = op(e3, e1)) | ~ sP46), inference(cnf_transformation, [], [f121])).
fof(f121, plain, ((~ (e2 = op(e4, e1)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e0, e1))) | ~ sP46), inference(nnf_transformation, [], [f53])).
fof(f53, plain, ((~ (e2 = op(e4, e1)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e0, e1))) | ~ sP46), inference(usedef, [], [e53])).
fof(e53, plain, (sP46 <=> (~ (e2 = op(e4, e1)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP46])])).
fof(f1477, plain, (sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1476, f925])).
fof(f925, plain, ~ sP47, inference(subsumption_resolution, [], [f366, f185])).
fof(f366, plain, (~ (e3 = op(e1, e2)) | ~ sP47), inference(cnf_transformation, [], [f120])).
fof(f120, plain, ((~ (e3 = op(e1, e4)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e1, e0))) | ~ sP47), inference(nnf_transformation, [], [f54])).
fof(f54, plain, ((~ (e3 = op(e1, e4)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e1, e0))) | ~ sP47), inference(usedef, [], [e54])).
fof(e54, plain, (sP47 <=> (~ (e3 = op(e1, e4)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP47])])).
fof(f1476, plain, (sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1475, f922])).
fof(f922, plain, ~ sP48, inference(subsumption_resolution, [], [f363, f199])).
fof(f363, plain, (~ (e3 = op(e4, e1)) | ~ sP48), inference(cnf_transformation, [], [f119])).
fof(f119, plain, ((~ (e3 = op(e4, e1)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e0, e1))) | ~ sP48), inference(nnf_transformation, [], [f55])).
fof(f55, plain, ((~ (e3 = op(e4, e1)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e0, e1))) | ~ sP48), inference(usedef, [], [e55])).
fof(e55, plain, (sP48 <=> (~ (e3 = op(e4, e1)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP48])])).
fof(f1475, plain, (sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1474, f921])).
fof(f921, plain, ~ sP49, inference(subsumption_resolution, [], [f355, f184])).
fof(f355, plain, (~ (e4 = op(e1, e1)) | ~ sP49), inference(cnf_transformation, [], [f118])).
fof(f118, plain, ((~ (e4 = op(e1, e4)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e1, e0))) | ~ sP49), inference(nnf_transformation, [], [f56])).
fof(f56, plain, ((~ (e4 = op(e1, e4)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e1, e0))) | ~ sP49), inference(usedef, [], [e56])).
fof(e56, plain, (sP49 <=> (~ (e4 = op(e1, e4)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP49])])).
fof(f1474, plain, (sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1473, f917])).
fof(f917, plain, ~ sP50, inference(subsumption_resolution, [], [f350, f184])).
fof(f350, plain, (~ (e4 = op(e1, e1)) | ~ sP50), inference(cnf_transformation, [], [f117])).
fof(f117, plain, ((~ (e4 = op(e4, e1)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e0, e1))) | ~ sP50), inference(nnf_transformation, [], [f57])).
fof(f57, plain, ((~ (e4 = op(e4, e1)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e0, e1))) | ~ sP50), inference(usedef, [], [e57])).
fof(e57, plain, (sP50 <=> (~ (e4 = op(e4, e1)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP50])])).
fof(f1473, plain, (sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1472, f913])).
fof(f913, plain, ~ sP51, inference(subsumption_resolution, [], [f677, f617])).
fof(f677, plain, (~ (op(e2, e1) = unit) | ~ sP51), inference(definition_unfolding, [], [f345, f203])).
fof(f345, plain, (~ (e0 = op(e2, e1)) | ~ sP51), inference(cnf_transformation, [], [f116])).
fof(f116, plain, ((~ (e0 = op(e2, e4)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e2, e0))) | ~ sP51), inference(nnf_transformation, [], [f58])).
fof(f58, plain, ((~ (e0 = op(e2, e4)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e2, e0))) | ~ sP51), inference(usedef, [], [e58])).
fof(e58, plain, (sP51 <=> (~ (e0 = op(e2, e4)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP51])])).
fof(f1472, plain, (sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1471, f909])).
fof(f909, plain, ~ sP52, inference(subsumption_resolution, [], [f670, f615])).
fof(f670, plain, (~ (op(e3, e2) = unit) | ~ sP52), inference(definition_unfolding, [], [f342, f203])).
fof(f342, plain, (~ (e0 = op(e3, e2)) | ~ sP52), inference(cnf_transformation, [], [f115])).
fof(f115, plain, ((~ (e0 = op(e4, e2)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e0, e2))) | ~ sP52), inference(nnf_transformation, [], [f59])).
fof(f59, plain, ((~ (e0 = op(e4, e2)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e0, e2))) | ~ sP52), inference(usedef, [], [e59])).
fof(e59, plain, (sP52 <=> (~ (e0 = op(e4, e2)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP52])])).
fof(f1471, plain, (sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1470, f907])).
fof(f907, plain, ~ sP53, inference(subsumption_resolution, [], [f337, f191])).
fof(f337, plain, (~ (e1 = op(e2, e3)) | ~ sP53), inference(cnf_transformation, [], [f114])).
fof(f114, plain, ((~ (e1 = op(e2, e4)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e2, e0))) | ~ sP53), inference(nnf_transformation, [], [f60])).
fof(f60, plain, ((~ (e1 = op(e2, e4)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e2, e0))) | ~ sP53), inference(usedef, [], [e60])).
fof(e60, plain, (sP53 <=> (~ (e1 = op(e2, e4)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP53])])).
fof(f1470, plain, (sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1469, f905])).
fof(f905, plain, ~ sP54, inference(subsumption_resolution, [], [f333, f200])).
fof(f333, plain, (~ (e1 = op(e4, e2)) | ~ sP54), inference(cnf_transformation, [], [f113])).
fof(f113, plain, ((~ (e1 = op(e4, e2)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e0, e2))) | ~ sP54), inference(nnf_transformation, [], [f61])).
fof(f61, plain, ((~ (e1 = op(e4, e2)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e0, e2))) | ~ sP54), inference(usedef, [], [e61])).
fof(e61, plain, (sP54 <=> (~ (e1 = op(e4, e2)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP54])])).
fof(f1469, plain, (sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1468, f904])).
fof(f904, plain, ~ sP55, inference(subsumption_resolution, [], [f666, f618])).
fof(f666, plain, (~ (e2 = op(e2, unit)) | ~ sP55), inference(definition_unfolding, [], [f324, f203])).
fof(f324, plain, (~ (e2 = op(e2, e0)) | ~ sP55), inference(cnf_transformation, [], [f112])).
fof(f112, plain, ((~ (e2 = op(e2, e4)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e2, e0))) | ~ sP55), inference(nnf_transformation, [], [f62])).
fof(f62, plain, ((~ (e2 = op(e2, e4)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e2, e0))) | ~ sP55), inference(usedef, [], [e62])).
fof(e62, plain, (sP55 <=> (~ (e2 = op(e2, e4)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP55])])).
fof(f1468, plain, (sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1467, f899])).
fof(f899, plain, ~ sP56, inference(subsumption_resolution, [], [f665, f623])).
fof(f665, plain, (~ (e2 = op(unit, e2)) | ~ sP56), inference(definition_unfolding, [], [f319, f203])).
fof(f319, plain, (~ (e2 = op(e0, e2)) | ~ sP56), inference(cnf_transformation, [], [f111])).
fof(f111, plain, ((~ (e2 = op(e4, e2)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e0, e2))) | ~ sP56), inference(nnf_transformation, [], [f63])).
fof(f63, plain, ((~ (e2 = op(e4, e2)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e0, e2))) | ~ sP56), inference(usedef, [], [e63])).
fof(e63, plain, (sP56 <=> (~ (e2 = op(e4, e2)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP56])])).
fof(f1467, plain, (sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1466, f894])).
fof(f894, plain, ~ sP57, inference(subsumption_resolution, [], [f318, f192])).
fof(f318, plain, (~ (e3 = op(e2, e4)) | ~ sP57), inference(cnf_transformation, [], [f110])).
fof(f110, plain, ((~ (e3 = op(e2, e4)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e2, e0))) | ~ sP57), inference(nnf_transformation, [], [f64])).
fof(f64, plain, ((~ (e3 = op(e2, e4)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e2, e0))) | ~ sP57), inference(usedef, [], [e64])).
fof(e64, plain, (sP57 <=> (~ (e3 = op(e2, e4)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP57])])).
fof(f1466, plain, (sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1465, f893])).
fof(f893, plain, ~ sP58, inference(subsumption_resolution, [], [f310, f185])).
fof(f310, plain, (~ (e3 = op(e1, e2)) | ~ sP58), inference(cnf_transformation, [], [f109])).
fof(f109, plain, ((~ (e3 = op(e4, e2)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e0, e2))) | ~ sP58), inference(nnf_transformation, [], [f65])).
fof(f65, plain, ((~ (e3 = op(e4, e2)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e0, e2))) | ~ sP58), inference(usedef, [], [e65])).
fof(e65, plain, (sP58 <=> (~ (e3 = op(e4, e2)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP58])])).
fof(f1465, plain, (sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1464, f889])).
fof(f889, plain, ~ sP59, inference(subsumption_resolution, [], [f306, f190])).
fof(f306, plain, (~ (e4 = op(e2, e2)) | ~ sP59), inference(cnf_transformation, [], [f108])).
fof(f108, plain, ((~ (e4 = op(e2, e4)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e2, e0))) | ~ sP59), inference(nnf_transformation, [], [f66])).
fof(f66, plain, ((~ (e4 = op(e2, e4)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e2, e0))) | ~ sP59), inference(usedef, [], [e66])).
fof(e66, plain, (sP59 <=> (~ (e4 = op(e2, e4)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP59])])).
fof(f1464, plain, (sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1463, f886])).
fof(f886, plain, ~ sP60, inference(subsumption_resolution, [], [f301, f190])).
fof(f301, plain, (~ (e4 = op(e2, e2)) | ~ sP60), inference(cnf_transformation, [], [f107])).
fof(f107, plain, ((~ (e4 = op(e4, e2)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e0, e2))) | ~ sP60), inference(nnf_transformation, [], [f67])).
fof(f67, plain, ((~ (e4 = op(e4, e2)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e0, e2))) | ~ sP60), inference(usedef, [], [e67])).
fof(e67, plain, (sP60 <=> (~ (e4 = op(e4, e2)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP60])])).
fof(f1463, plain, (sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1462, f883])).
fof(f883, plain, ~ sP61, inference(subsumption_resolution, [], [f658, f615])).
fof(f658, plain, (~ (op(e3, e2) = unit) | ~ sP61), inference(definition_unfolding, [], [f296, f203])).
fof(f296, plain, (~ (e0 = op(e3, e2)) | ~ sP61), inference(cnf_transformation, [], [f106])).
fof(f106, plain, ((~ (e0 = op(e3, e4)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e3, e0))) | ~ sP61), inference(nnf_transformation, [], [f68])).
fof(f68, plain, ((~ (e0 = op(e3, e4)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e3, e0))) | ~ sP61), inference(usedef, [], [e68])).
fof(e68, plain, (sP61 <=> (~ (e0 = op(e3, e4)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP61])])).
fof(f1462, plain, (sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1461, f880])).
fof(f880, plain, ~ sP62, inference(subsumption_resolution, [], [f654, f619])).
fof(f654, plain, (~ (op(e1, e3) = unit) | ~ sP62), inference(definition_unfolding, [], [f290, f203])).
fof(f290, plain, (~ (e0 = op(e1, e3)) | ~ sP62), inference(cnf_transformation, [], [f105])).
fof(f105, plain, ((~ (e0 = op(e4, e3)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e0, e3))) | ~ sP62), inference(nnf_transformation, [], [f69])).
fof(f69, plain, ((~ (e0 = op(e4, e3)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e0, e3))) | ~ sP62), inference(usedef, [], [e69])).
fof(e69, plain, (sP62 <=> (~ (e0 = op(e4, e3)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP62])])).
fof(f1461, plain, (sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1460, f876])).
fof(f876, plain, ~ sP63, inference(subsumption_resolution, [], [f288, f197])).
fof(f288, plain, (~ (e1 = op(e3, e4)) | ~ sP63), inference(cnf_transformation, [], [f104])).
fof(f104, plain, ((~ (e1 = op(e3, e4)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e3, e0))) | ~ sP63), inference(nnf_transformation, [], [f70])).
fof(f70, plain, ((~ (e1 = op(e3, e4)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e3, e0))) | ~ sP63), inference(usedef, [], [e70])).
fof(e70, plain, (sP63 <=> (~ (e1 = op(e3, e4)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP63])])).
fof(f1460, plain, (sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1459, f875])).
fof(f875, plain, ~ sP64, inference(subsumption_resolution, [], [f281, f191])).
fof(f281, plain, (~ (e1 = op(e2, e3)) | ~ sP64), inference(cnf_transformation, [], [f103])).
fof(f103, plain, ((~ (e1 = op(e4, e3)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e0, e3))) | ~ sP64), inference(nnf_transformation, [], [f71])).
fof(f71, plain, ((~ (e1 = op(e4, e3)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e0, e3))) | ~ sP64), inference(usedef, [], [e71])).
fof(e71, plain, (sP64 <=> (~ (e1 = op(e4, e3)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP64])])).
fof(f1459, plain, (sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1458, f872])).
fof(f872, plain, ~ sP65, inference(subsumption_resolution, [], [f275, f194])).
fof(f275, plain, (~ (e2 = op(e3, e1)) | ~ sP65), inference(cnf_transformation, [], [f102])).
fof(f102, plain, ((~ (e2 = op(e3, e4)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e3, e0))) | ~ sP65), inference(nnf_transformation, [], [f72])).
fof(f72, plain, ((~ (e2 = op(e3, e4)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e3, e0))) | ~ sP65), inference(usedef, [], [e72])).
fof(e72, plain, (sP65 <=> (~ (e2 = op(e3, e4)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP65])])).
fof(f1458, plain, (sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1457, f868])).
fof(f868, plain, ~ sP66, inference(subsumption_resolution, [], [f273, f201])).
fof(f273, plain, (~ (e2 = op(e4, e3)) | ~ sP66), inference(cnf_transformation, [], [f101])).
fof(f101, plain, ((~ (e2 = op(e4, e3)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e0, e3))) | ~ sP66), inference(nnf_transformation, [], [f73])).
fof(f73, plain, ((~ (e2 = op(e4, e3)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e0, e3))) | ~ sP66), inference(usedef, [], [e73])).
fof(e73, plain, (sP66 <=> (~ (e2 = op(e4, e3)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP66])])).
fof(f1457, plain, (sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1456, f867])).
fof(f867, plain, ~ sP67, inference(subsumption_resolution, [], [f646, f616])).
fof(f646, plain, (~ (e3 = op(e3, unit)) | ~ sP67), inference(definition_unfolding, [], [f264, f203])).
fof(f264, plain, (~ (e3 = op(e3, e0)) | ~ sP67), inference(cnf_transformation, [], [f100])).
fof(f100, plain, ((~ (e3 = op(e3, e4)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e3, e0))) | ~ sP67), inference(nnf_transformation, [], [f74])).
fof(f74, plain, ((~ (e3 = op(e3, e4)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e3, e0))) | ~ sP67), inference(usedef, [], [e74])).
fof(e74, plain, (sP67 <=> (~ (e3 = op(e3, e4)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP67])])).
fof(f1456, plain, (sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1455, f862])).
fof(f862, plain, ~ sP68, inference(subsumption_resolution, [], [f645, f622])).
fof(f645, plain, (~ (e3 = op(unit, e3)) | ~ sP68), inference(definition_unfolding, [], [f259, f203])).
fof(f259, plain, (~ (e3 = op(e0, e3)) | ~ sP68), inference(cnf_transformation, [], [f99])).
fof(f99, plain, ((~ (e3 = op(e4, e3)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e0, e3))) | ~ sP68), inference(nnf_transformation, [], [f75])).
fof(f75, plain, ((~ (e3 = op(e4, e3)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e0, e3))) | ~ sP68), inference(usedef, [], [e75])).
fof(e75, plain, (sP68 <=> (~ (e3 = op(e4, e3)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP68])])).
fof(f1455, plain, (sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1454, f857])).
fof(f857, plain, ~ sP69, inference(subsumption_resolution, [], [f257, f196])).
fof(f257, plain, (~ (e4 = op(e3, e3)) | ~ sP69), inference(cnf_transformation, [], [f98])).
fof(f98, plain, ((~ (e4 = op(e3, e4)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e3, e0))) | ~ sP69), inference(nnf_transformation, [], [f76])).
fof(f76, plain, ((~ (e4 = op(e3, e4)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e3, e0))) | ~ sP69), inference(usedef, [], [e76])).
fof(e76, plain, (sP69 <=> (~ (e4 = op(e3, e4)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP69])])).
fof(f1454, plain, (sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1453, f855])).
fof(f855, plain, ~ sP70, inference(subsumption_resolution, [], [f252, f196])).
fof(f252, plain, (~ (e4 = op(e3, e3)) | ~ sP70), inference(cnf_transformation, [], [f97])).
fof(f97, plain, ((~ (e4 = op(e4, e3)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e0, e3))) | ~ sP70), inference(nnf_transformation, [], [f77])).
fof(f77, plain, ((~ (e4 = op(e4, e3)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e0, e3))) | ~ sP70), inference(usedef, [], [e77])).
fof(e77, plain, (sP70 <=> (~ (e4 = op(e4, e3)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP70])])).
fof(f1453, plain, (sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1452, f853])).
fof(f853, plain, ~ sP71, inference(subsumption_resolution, [], [f638, f613])).
fof(f638, plain, (~ (op(e4, e4) = unit) | ~ sP71), inference(definition_unfolding, [], [f248, f203])).
fof(f248, plain, (~ (e0 = op(e4, e4)) | ~ sP71), inference(cnf_transformation, [], [f96])).
fof(f96, plain, ((~ (e0 = op(e4, e4)) & ~ (e0 = op(e4, e3)) & ~ (e0 = op(e4, e2)) & ~ (e0 = op(e4, e1)) & ~ (e0 = op(e4, e0))) | ~ sP71), inference(nnf_transformation, [], [f78])).
fof(f78, plain, ((~ (e0 = op(e4, e4)) & ~ (e0 = op(e4, e3)) & ~ (e0 = op(e4, e2)) & ~ (e0 = op(e4, e1)) & ~ (e0 = op(e4, e0))) | ~ sP71), inference(usedef, [], [e78])).
fof(e78, plain, (sP71 <=> (~ (e0 = op(e4, e4)) & ~ (e0 = op(e4, e3)) & ~ (e0 = op(e4, e2)) & ~ (e0 = op(e4, e1)) & ~ (e0 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP71])])).
fof(f1452, plain, (sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1451, f852])).
fof(f852, plain, ~ sP72, inference(subsumption_resolution, [], [f633, f613])).
fof(f633, plain, (~ (op(e4, e4) = unit) | ~ sP72), inference(definition_unfolding, [], [f243, f203])).
fof(f243, plain, (~ (e0 = op(e4, e4)) | ~ sP72), inference(cnf_transformation, [], [f95])).
fof(f95, plain, ((~ (e0 = op(e4, e4)) & ~ (e0 = op(e3, e4)) & ~ (e0 = op(e2, e4)) & ~ (e0 = op(e1, e4)) & ~ (e0 = op(e0, e4))) | ~ sP72), inference(nnf_transformation, [], [f79])).
fof(f79, plain, ((~ (e0 = op(e4, e4)) & ~ (e0 = op(e3, e4)) & ~ (e0 = op(e2, e4)) & ~ (e0 = op(e1, e4)) & ~ (e0 = op(e0, e4))) | ~ sP72), inference(usedef, [], [e79])).
fof(e79, plain, (sP72 <=> (~ (e0 = op(e4, e4)) & ~ (e0 = op(e3, e4)) & ~ (e0 = op(e2, e4)) & ~ (e0 = op(e1, e4)) & ~ (e0 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP72])])).
fof(f1451, plain, (sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1450, f851])).
fof(f851, plain, ~ sP73, inference(subsumption_resolution, [], [f236, f200])).
fof(f236, plain, (~ (e1 = op(e4, e2)) | ~ sP73), inference(cnf_transformation, [], [f94])).
fof(f94, plain, ((~ (e1 = op(e4, e4)) & ~ (e1 = op(e4, e3)) & ~ (e1 = op(e4, e2)) & ~ (e1 = op(e4, e1)) & ~ (e1 = op(e4, e0))) | ~ sP73), inference(nnf_transformation, [], [f80])).
fof(f80, plain, ((~ (e1 = op(e4, e4)) & ~ (e1 = op(e4, e3)) & ~ (e1 = op(e4, e2)) & ~ (e1 = op(e4, e1)) & ~ (e1 = op(e4, e0))) | ~ sP73), inference(usedef, [], [e80])).
fof(e80, plain, (sP73 <=> (~ (e1 = op(e4, e4)) & ~ (e1 = op(e4, e3)) & ~ (e1 = op(e4, e2)) & ~ (e1 = op(e4, e1)) & ~ (e1 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP73])])).
fof(f1450, plain, (sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1449, f848])).
fof(f848, plain, ~ sP74, inference(subsumption_resolution, [], [f232, f197])).
fof(f232, plain, (~ (e1 = op(e3, e4)) | ~ sP74), inference(cnf_transformation, [], [f93])).
fof(f93, plain, ((~ (e1 = op(e4, e4)) & ~ (e1 = op(e3, e4)) & ~ (e1 = op(e2, e4)) & ~ (e1 = op(e1, e4)) & ~ (e1 = op(e0, e4))) | ~ sP74), inference(nnf_transformation, [], [f81])).
fof(f81, plain, ((~ (e1 = op(e4, e4)) & ~ (e1 = op(e3, e4)) & ~ (e1 = op(e2, e4)) & ~ (e1 = op(e1, e4)) & ~ (e1 = op(e0, e4))) | ~ sP74), inference(usedef, [], [e81])).
fof(e81, plain, (sP74 <=> (~ (e1 = op(e4, e4)) & ~ (e1 = op(e3, e4)) & ~ (e1 = op(e2, e4)) & ~ (e1 = op(e1, e4)) & ~ (e1 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP74])])).
fof(f1449, plain, (sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1448, f846])).
fof(f846, plain, ~ sP75, inference(subsumption_resolution, [], [f227, f201])).
fof(f227, plain, (~ (e2 = op(e4, e3)) | ~ sP75), inference(cnf_transformation, [], [f92])).
fof(f92, plain, ((~ (e2 = op(e4, e4)) & ~ (e2 = op(e4, e3)) & ~ (e2 = op(e4, e2)) & ~ (e2 = op(e4, e1)) & ~ (e2 = op(e4, e0))) | ~ sP75), inference(nnf_transformation, [], [f82])).
fof(f82, plain, ((~ (e2 = op(e4, e4)) & ~ (e2 = op(e4, e3)) & ~ (e2 = op(e4, e2)) & ~ (e2 = op(e4, e1)) & ~ (e2 = op(e4, e0))) | ~ sP75), inference(usedef, [], [e82])).
fof(e82, plain, (sP75 <=> (~ (e2 = op(e4, e4)) & ~ (e2 = op(e4, e3)) & ~ (e2 = op(e4, e2)) & ~ (e2 = op(e4, e1)) & ~ (e2 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP75])])).
fof(f1448, plain, (sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1447, f844])).
fof(f844, plain, ~ sP76, inference(subsumption_resolution, [], [f220, f187])).
fof(f220, plain, (~ (e2 = op(e1, e4)) | ~ sP76), inference(cnf_transformation, [], [f91])).
fof(f91, plain, ((~ (e2 = op(e4, e4)) & ~ (e2 = op(e3, e4)) & ~ (e2 = op(e2, e4)) & ~ (e2 = op(e1, e4)) & ~ (e2 = op(e0, e4))) | ~ sP76), inference(nnf_transformation, [], [f83])).
fof(f83, plain, ((~ (e2 = op(e4, e4)) & ~ (e2 = op(e3, e4)) & ~ (e2 = op(e2, e4)) & ~ (e2 = op(e1, e4)) & ~ (e2 = op(e0, e4))) | ~ sP76), inference(usedef, [], [e83])).
fof(e83, plain, (sP76 <=> (~ (e2 = op(e4, e4)) & ~ (e2 = op(e3, e4)) & ~ (e2 = op(e2, e4)) & ~ (e2 = op(e1, e4)) & ~ (e2 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP76])])).
fof(f1447, plain, (sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1446, f840])).
fof(f840, plain, ~ sP77, inference(subsumption_resolution, [], [f215, f199])).
fof(f215, plain, (~ (e3 = op(e4, e1)) | ~ sP77), inference(cnf_transformation, [], [f90])).
fof(f90, plain, ((~ (e3 = op(e4, e4)) & ~ (e3 = op(e4, e3)) & ~ (e3 = op(e4, e2)) & ~ (e3 = op(e4, e1)) & ~ (e3 = op(e4, e0))) | ~ sP77), inference(nnf_transformation, [], [f84])).
fof(f84, plain, ((~ (e3 = op(e4, e4)) & ~ (e3 = op(e4, e3)) & ~ (e3 = op(e4, e2)) & ~ (e3 = op(e4, e1)) & ~ (e3 = op(e4, e0))) | ~ sP77), inference(usedef, [], [e84])).
fof(e84, plain, (sP77 <=> (~ (e3 = op(e4, e4)) & ~ (e3 = op(e4, e3)) & ~ (e3 = op(e4, e2)) & ~ (e3 = op(e4, e1)) & ~ (e3 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP77])])).
fof(f1446, plain, (sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1445, f836])).
fof(f836, plain, ~ sP78, inference(subsumption_resolution, [], [f211, f192])).
fof(f211, plain, (~ (e3 = op(e2, e4)) | ~ sP78), inference(cnf_transformation, [], [f89])).
fof(f89, plain, ((~ (e3 = op(e4, e4)) & ~ (e3 = op(e3, e4)) & ~ (e3 = op(e2, e4)) & ~ (e3 = op(e1, e4)) & ~ (e3 = op(e0, e4))) | ~ sP78), inference(nnf_transformation, [], [f85])).
fof(f85, plain, ((~ (e3 = op(e4, e4)) & ~ (e3 = op(e3, e4)) & ~ (e3 = op(e2, e4)) & ~ (e3 = op(e1, e4)) & ~ (e3 = op(e0, e4))) | ~ sP78), inference(usedef, [], [e85])).
fof(e85, plain, (sP78 <=> (~ (e3 = op(e4, e4)) & ~ (e3 = op(e3, e4)) & ~ (e3 = op(e2, e4)) & ~ (e3 = op(e1, e4)) & ~ (e3 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP78])])).
fof(f1445, plain, (sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1444, f833])).
fof(f833, plain, ~ sP79, inference(subsumption_resolution, [], [f626, f614])).
fof(f626, plain, (~ (e4 = op(e4, unit)) | ~ sP79), inference(definition_unfolding, [], [f204, f203])).
fof(f204, plain, (~ (e4 = op(e4, e0)) | ~ sP79), inference(cnf_transformation, [], [f88])).
fof(f88, plain, ((~ (e4 = op(e4, e4)) & ~ (e4 = op(e4, e3)) & ~ (e4 = op(e4, e2)) & ~ (e4 = op(e4, e1)) & ~ (e4 = op(e4, e0))) | ~ sP79), inference(nnf_transformation, [], [f86])).
fof(f86, plain, ((~ (e4 = op(e4, e4)) & ~ (e4 = op(e4, e3)) & ~ (e4 = op(e4, e2)) & ~ (e4 = op(e4, e1)) & ~ (e4 = op(e4, e0))) | ~ sP79), inference(usedef, [], [e86])).
fof(e86, plain, (sP79 <=> (~ (e4 = op(e4, e4)) & ~ (e4 = op(e4, e3)) & ~ (e4 = op(e4, e2)) & ~ (e4 = op(e4, e1)) & ~ (e4 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP79])])).
fof(f1444, plain, (sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f823, f621])).
fof(f823, plain, (~ (e4 = op(unit, e4)) | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (e4 = op(e4, unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(duplicate_literal_removal, [], [f822])).
fof(f822, plain, (~ (e4 = op(unit, e4)) | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | ~ (unit = op(unit, unit)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(definition_unfolding, [], [f604, f203, f203, f203, f203, f203])).
fof(f604, plain, (~ (e4 = op(e0, e4)) | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (e0 = op(e0, unit)) | ~ (e0 = op(unit, e0)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f87])).
fof(f87, plain, ((~ (e4 = op(e4, e4)) & ~ (e4 = op(e3, e4)) & ~ (e4 = op(e2, e4)) & ~ (e4 = op(e1, e4)) & ~ (e4 = op(e0, e4))) | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (e0 = op(e0, unit)) | ~ (e0 = op(unit, e0)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(definition_folding, [], [f6, e86, e85, e84, e83, e82, e81, e80, e79, e78, e77, e76, e75, e74, e73, e72, e71, e70, e69, e68, e67, e66, e65, e64, e63, e62, e61, e60, e59, e58, e57, e56, e55, e54, e53, e52, e51, e50, e49, e48, e47, e46, e45, e44, e43, e42, e41, e40, e39, e38, e37, e36, e35, e34, e33, e32, e31, e30, e29, e28, e27, e26, e25, e24, e23, e22, e21, e20, e19, e18, e17, e16, e15, e14, e13, e12, e11, e10, e9, e8, e7])).
fof(f6, plain, ((~ (e4 = op(e4, e4)) & ~ (e4 = op(e3, e4)) & ~ (e4 = op(e2, e4)) & ~ (e4 = op(e1, e4)) & ~ (e4 = op(e0, e4))) | (~ (e4 = op(e4, e4)) & ~ (e4 = op(e4, e3)) & ~ (e4 = op(e4, e2)) & ~ (e4 = op(e4, e1)) & ~ (e4 = op(e4, e0))) | (~ (e3 = op(e4, e4)) & ~ (e3 = op(e3, e4)) & ~ (e3 = op(e2, e4)) & ~ (e3 = op(e1, e4)) & ~ (e3 = op(e0, e4))) | (~ (e3 = op(e4, e4)) & ~ (e3 = op(e4, e3)) & ~ (e3 = op(e4, e2)) & ~ (e3 = op(e4, e1)) & ~ (e3 = op(e4, e0))) | (~ (e2 = op(e4, e4)) & ~ (e2 = op(e3, e4)) & ~ (e2 = op(e2, e4)) & ~ (e2 = op(e1, e4)) & ~ (e2 = op(e0, e4))) | (~ (e2 = op(e4, e4)) & ~ (e2 = op(e4, e3)) & ~ (e2 = op(e4, e2)) & ~ (e2 = op(e4, e1)) & ~ (e2 = op(e4, e0))) | (~ (e1 = op(e4, e4)) & ~ (e1 = op(e3, e4)) & ~ (e1 = op(e2, e4)) & ~ (e1 = op(e1, e4)) & ~ (e1 = op(e0, e4))) | (~ (e1 = op(e4, e4)) & ~ (e1 = op(e4, e3)) & ~ (e1 = op(e4, e2)) & ~ (e1 = op(e4, e1)) & ~ (e1 = op(e4, e0))) | (~ (e0 = op(e4, e4)) & ~ (e0 = op(e3, e4)) & ~ (e0 = op(e2, e4)) & ~ (e0 = op(e1, e4)) & ~ (e0 = op(e0, e4))) | (~ (e0 = op(e4, e4)) & ~ (e0 = op(e4, e3)) & ~ (e0 = op(e4, e2)) & ~ (e0 = op(e4, e1)) & ~ (e0 = op(e4, e0))) | (~ (e4 = op(e4, e3)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e0, e3))) | (~ (e4 = op(e3, e4)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e3, e0))) | (~ (e3 = op(e4, e3)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e0, e3))) | (~ (e3 = op(e3, e4)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e3, e0))) | (~ (e2 = op(e4, e3)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e0, e3))) | (~ (e2 = op(e3, e4)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e3, e0))) | (~ (e1 = op(e4, e3)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e0, e3))) | (~ (e1 = op(e3, e4)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e3, e0))) | (~ (e0 = op(e4, e3)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e0, e3))) | (~ (e0 = op(e3, e4)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e3, e0))) | (~ (e4 = op(e4, e2)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e0, e2))) | (~ (e4 = op(e2, e4)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e2, e0))) | (~ (e3 = op(e4, e2)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e0, e2))) | (~ (e3 = op(e2, e4)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e2, e0))) | (~ (e2 = op(e4, e2)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e0, e2))) | (~ (e2 = op(e2, e4)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e2, e0))) | (~ (e1 = op(e4, e2)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e0, e2))) | (~ (e1 = op(e2, e4)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e2, e0))) | (~ (e0 = op(e4, e2)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e0, e2))) | (~ (e0 = op(e2, e4)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e2, e0))) | (~ (e4 = op(e4, e1)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e0, e1))) | (~ (e4 = op(e1, e4)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e1, e0))) | (~ (e3 = op(e4, e1)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e0, e1))) | (~ (e3 = op(e1, e4)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e1, e0))) | (~ (e2 = op(e4, e1)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e0, e1))) | (~ (e2 = op(e1, e4)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e1, e0))) | (~ (e1 = op(e4, e1)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e0, e1))) | (~ (e1 = op(e1, e4)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e1, e0))) | (~ (e0 = op(e4, e1)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e0, e1))) | (~ (e0 = op(e1, e4)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e1, e0))) | (~ (e4 = op(e4, e0)) & ~ (e4 = op(e3, e0)) & ~ (e4 = op(e2, e0)) & ~ (e4 = op(e1, e0)) & ~ (e4 = op(e0, e0))) | (~ (e4 = op(e0, e4)) & ~ (e4 = op(e0, e3)) & ~ (e4 = op(e0, e2)) & ~ (e4 = op(e0, e1)) & ~ (e4 = op(e0, e0))) | (~ (e3 = op(e4, e0)) & ~ (e3 = op(e3, e0)) & ~ (e3 = op(e2, e0)) & ~ (e3 = op(e1, e0)) & ~ (e3 = op(e0, e0))) | (~ (e3 = op(e0, e4)) & ~ (e3 = op(e0, e3)) & ~ (e3 = op(e0, e2)) & ~ (e3 = op(e0, e1)) & ~ (e3 = op(e0, e0))) | (~ (e2 = op(e4, e0)) & ~ (e2 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e2 = op(e1, e0)) & ~ (e2 = op(e0, e0))) | (~ (e2 = op(e0, e4)) & ~ (e2 = op(e0, e3)) & ~ (e2 = op(e0, e2)) & ~ (e2 = op(e0, e1)) & ~ (e2 = op(e0, e0))) | (~ (e1 = op(e4, e0)) & ~ (e1 = op(e3, e0)) & ~ (e1 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e1 = op(e0, e0))) | (~ (e1 = op(e0, e4)) & ~ (e1 = op(e0, e3)) & ~ (e1 = op(e0, e2)) & ~ (e1 = op(e0, e1)) & ~ (e1 = op(e0, e0))) | (~ (e0 = op(e4, e0)) & ~ (e0 = op(e3, e0)) & ~ (e0 = op(e2, e0)) & ~ (e0 = op(e1, e0)) & ~ (e0 = op(e0, e0))) | (~ (e0 = op(e0, e4)) & ~ (e0 = op(e0, e3)) & ~ (e0 = op(e0, e2)) & ~ (e0 = op(e0, e1)) & ~ (e0 = op(e0, e0))) | (~ (e4 = unit) & ~ (e3 = unit) & ~ (e2 = unit) & ~ (e1 = unit) & ~ (e0 = unit)) | ~ (e4 = op(e4, unit)) | ~ (e4 = op(unit, e4)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (e0 = op(e0, unit)) | ~ (e0 = op(unit, e0)) | (~ (e4 = op(e4, e4)) & ~ (e3 = op(e4, e4)) & ~ (e2 = op(e4, e4)) & ~ (e1 = op(e4, e4)) & ~ (e0 = op(e4, e4))) | (~ (e4 = op(e4, e3)) & ~ (e3 = op(e4, e3)) & ~ (e2 = op(e4, e3)) & ~ (e1 = op(e4, e3)) & ~ (e0 = op(e4, e3))) | (~ (e4 = op(e4, e2)) & ~ (e3 = op(e4, e2)) & ~ (e2 = op(e4, e2)) & ~ (e1 = op(e4, e2)) & ~ (e0 = op(e4, e2))) | (~ (e4 = op(e4, e1)) & ~ (e3 = op(e4, e1)) & ~ (e2 = op(e4, e1)) & ~ (e1 = op(e4, e1)) & ~ (e0 = op(e4, e1))) | (~ (e4 = op(e4, e0)) & ~ (e3 = op(e4, e0)) & ~ (e2 = op(e4, e0)) & ~ (e1 = op(e4, e0)) & ~ (e0 = op(e4, e0))) | (~ (e4 = op(e3, e4)) & ~ (e3 = op(e3, e4)) & ~ (e2 = op(e3, e4)) & ~ (e1 = op(e3, e4)) & ~ (e0 = op(e3, e4))) | (~ (e4 = op(e3, e3)) & ~ (e3 = op(e3, e3)) & ~ (e2 = op(e3, e3)) & ~ (e1 = op(e3, e3)) & ~ (e0 = op(e3, e3))) | (~ (e4 = op(e3, e2)) & ~ (e3 = op(e3, e2)) & ~ (e2 = op(e3, e2)) & ~ (e1 = op(e3, e2)) & ~ (e0 = op(e3, e2))) | (~ (e4 = op(e3, e1)) & ~ (e3 = op(e3, e1)) & ~ (e2 = op(e3, e1)) & ~ (e1 = op(e3, e1)) & ~ (e0 = op(e3, e1))) | (~ (e4 = op(e3, e0)) & ~ (e3 = op(e3, e0)) & ~ (e2 = op(e3, e0)) & ~ (e1 = op(e3, e0)) & ~ (e0 = op(e3, e0))) | (~ (e4 = op(e2, e4)) & ~ (e3 = op(e2, e4)) & ~ (e2 = op(e2, e4)) & ~ (e1 = op(e2, e4)) & ~ (e0 = op(e2, e4))) | (~ (e4 = op(e2, e3)) & ~ (e3 = op(e2, e3)) & ~ (e2 = op(e2, e3)) & ~ (e1 = op(e2, e3)) & ~ (e0 = op(e2, e3))) | (~ (e4 = op(e2, e2)) & ~ (e3 = op(e2, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e2, e2)) & ~ (e0 = op(e2, e2))) | (~ (e4 = op(e2, e1)) & ~ (e3 = op(e2, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e2, e1)) & ~ (e0 = op(e2, e1))) | (~ (e4 = op(e2, e0)) & ~ (e3 = op(e2, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e2, e0)) & ~ (e0 = op(e2, e0))) | (~ (e4 = op(e1, e4)) & ~ (e3 = op(e1, e4)) & ~ (e2 = op(e1, e4)) & ~ (e1 = op(e1, e4)) & ~ (e0 = op(e1, e4))) | (~ (e4 = op(e1, e3)) & ~ (e3 = op(e1, e3)) & ~ (e2 = op(e1, e3)) & ~ (e1 = op(e1, e3)) & ~ (e0 = op(e1, e3))) | (~ (e4 = op(e1, e2)) & ~ (e3 = op(e1, e2)) & ~ (e2 = op(e1, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e1, e2))) | (~ (e4 = op(e1, e1)) & ~ (e3 = op(e1, e1)) & ~ (e2 = op(e1, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e1, e1))) | (~ (e4 = op(e1, e0)) & ~ (e3 = op(e1, e0)) & ~ (e2 = op(e1, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e1, e0))) | (~ (e4 = op(e0, e4)) & ~ (e3 = op(e0, e4)) & ~ (e2 = op(e0, e4)) & ~ (e1 = op(e0, e4)) & ~ (e0 = op(e0, e4))) | (~ (e4 = op(e0, e3)) & ~ (e3 = op(e0, e3)) & ~ (e2 = op(e0, e3)) & ~ (e1 = op(e0, e3)) & ~ (e0 = op(e0, e3))) | (~ (e4 = op(e0, e2)) & ~ (e3 = op(e0, e2)) & ~ (e2 = op(e0, e2)) & ~ (e1 = op(e0, e2)) & ~ (e0 = op(e0, e2))) | (~ (e4 = op(e0, e1)) & ~ (e3 = op(e0, e1)) & ~ (e2 = op(e0, e1)) & ~ (e1 = op(e0, e1)) & ~ (e0 = op(e0, e1))) | (~ (e4 = op(e0, e0)) & ~ (e3 = op(e0, e0)) & ~ (e2 = op(e0, e0)) & ~ (e1 = op(e0, e0)) & ~ (e0 = op(e0, e0))) | ((~ (e4 = op(e4, op(e4, e4))) | (e4 = op(e4, op(e4, e4)))) & (~ (e3 = op(e4, op(e3, e4))) | (e4 = op(e3, op(e3, e4)))) & (~ (e2 = op(e4, op(e2, e4))) | (e4 = op(e2, op(e2, e4)))) & (~ (e1 = op(e4, op(e1, e4))) | (e4 = op(e1, op(e1, e4)))) & (~ (e0 = op(e4, op(e0, e4))) | (e4 = op(e0, op(e0, e4))))) | ((~ (e4 = op(e3, op(e4, e3))) | (e3 = op(e4, op(e4, e3)))) & (~ (e3 = op(e3, op(e3, e3))) | (e3 = op(e3, op(e3, e3)))) & (~ (e2 = op(e3, op(e2, e3))) | (e3 = op(e2, op(e2, e3)))) & (~ (e1 = op(e3, op(e1, e3))) | (e3 = op(e1, op(e1, e3)))) & (~ (e0 = op(e3, op(e0, e3))) | (e3 = op(e0, op(e0, e3))))) | ((~ (e4 = op(e2, op(e4, e2))) | (e2 = op(e4, op(e4, e2)))) & (~ (e3 = op(e2, op(e3, e2))) | (e2 = op(e3, op(e3, e2)))) & (~ (e2 = op(e2, op(e2, e2))) | (e2 = op(e2, op(e2, e2)))) & (~ (e1 = op(e2, op(e1, e2))) | (e2 = op(e1, op(e1, e2)))) & (~ (e0 = op(e2, op(e0, e2))) | (e2 = op(e0, op(e0, e2))))) | ((~ (e4 = op(e1, op(e4, e1))) | (e1 = op(e4, op(e4, e1)))) & (~ (e3 = op(e1, op(e3, e1))) | (e1 = op(e3, op(e3, e1)))) & (~ (e2 = op(e1, op(e2, e1))) | (e1 = op(e2, op(e2, e1)))) & (~ (e1 = op(e1, op(e1, e1))) | (e1 = op(e1, op(e1, e1)))) & (~ (e0 = op(e1, op(e0, e1))) | (e1 = op(e0, op(e0, e1))))) | ((~ (e4 = op(e0, op(e4, e0))) | (e0 = op(e4, op(e4, e0)))) & (~ (e3 = op(e0, op(e3, e0))) | (e0 = op(e3, op(e3, e0)))) & (~ (e2 = op(e0, op(e2, e0))) | (e0 = op(e2, op(e2, e0)))) & (~ (e1 = op(e0, op(e1, e0))) | (e0 = op(e1, op(e1, e0)))) & (~ (e0 = op(e0, op(e0, e0))) | (e0 = op(e0, op(e0, e0)))))), inference(ennf_transformation, [], [f5])).
fof(f5, plain, ~ (((e4 = op(e4, e4)) | (e4 = op(e3, e4)) | (e4 = op(e2, e4)) | (e4 = op(e1, e4)) | (e4 = op(e0, e4))) & ((e4 = op(e4, e4)) | (e4 = op(e4, e3)) | (e4 = op(e4, e2)) | (e4 = op(e4, e1)) | (e4 = op(e4, e0))) & ((e3 = op(e4, e4)) | (e3 = op(e3, e4)) | (e3 = op(e2, e4)) | (e3 = op(e1, e4)) | (e3 = op(e0, e4))) & ((e3 = op(e4, e4)) | (e3 = op(e4, e3)) | (e3 = op(e4, e2)) | (e3 = op(e4, e1)) | (e3 = op(e4, e0))) & ((e2 = op(e4, e4)) | (e2 = op(e3, e4)) | (e2 = op(e2, e4)) | (e2 = op(e1, e4)) | (e2 = op(e0, e4))) & ((e2 = op(e4, e4)) | (e2 = op(e4, e3)) | (e2 = op(e4, e2)) | (e2 = op(e4, e1)) | (e2 = op(e4, e0))) & ((e1 = op(e4, e4)) | (e1 = op(e3, e4)) | (e1 = op(e2, e4)) | (e1 = op(e1, e4)) | (e1 = op(e0, e4))) & ((e1 = op(e4, e4)) | (e1 = op(e4, e3)) | (e1 = op(e4, e2)) | (e1 = op(e4, e1)) | (e1 = op(e4, e0))) & ((e0 = op(e4, e4)) | (e0 = op(e3, e4)) | (e0 = op(e2, e4)) | (e0 = op(e1, e4)) | (e0 = op(e0, e4))) & ((e0 = op(e4, e4)) | (e0 = op(e4, e3)) | (e0 = op(e4, e2)) | (e0 = op(e4, e1)) | (e0 = op(e4, e0))) & ((e4 = op(e4, e3)) | (e4 = op(e3, e3)) | (e4 = op(e2, e3)) | (e4 = op(e1, e3)) | (e4 = op(e0, e3))) & ((e4 = op(e3, e4)) | (e4 = op(e3, e3)) | (e4 = op(e3, e2)) | (e4 = op(e3, e1)) | (e4 = op(e3, e0))) & ((e3 = op(e4, e3)) | (e3 = op(e3, e3)) | (e3 = op(e2, e3)) | (e3 = op(e1, e3)) | (e3 = op(e0, e3))) & ((e3 = op(e3, e4)) | (e3 = op(e3, e3)) | (e3 = op(e3, e2)) | (e3 = op(e3, e1)) | (e3 = op(e3, e0))) & ((e2 = op(e4, e3)) | (e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(e0, e3))) & ((e2 = op(e3, e4)) | (e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, e0))) & ((e1 = op(e4, e3)) | (e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(e0, e3))) & ((e1 = op(e3, e4)) | (e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, e0))) & ((e0 = op(e4, e3)) | (e0 = op(e3, e3)) | (e0 = op(e2, e3)) | (e0 = op(e1, e3)) | (e0 = op(e0, e3))) & ((e0 = op(e3, e4)) | (e0 = op(e3, e3)) | (e0 = op(e3, e2)) | (e0 = op(e3, e1)) | (e0 = op(e3, e0))) & ((e4 = op(e4, e2)) | (e4 = op(e3, e2)) | (e4 = op(e2, e2)) | (e4 = op(e1, e2)) | (e4 = op(e0, e2))) & ((e4 = op(e2, e4)) | (e4 = op(e2, e3)) | (e4 = op(e2, e2)) | (e4 = op(e2, e1)) | (e4 = op(e2, e0))) & ((e3 = op(e4, e2)) | (e3 = op(e3, e2)) | (e3 = op(e2, e2)) | (e3 = op(e1, e2)) | (e3 = op(e0, e2))) & ((e3 = op(e2, e4)) | (e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, e0))) & ((e2 = op(e4, e2)) | (e2 = op(e3, e2)) | (e2 = op(e2, e2)) | (e2 = op(e1, e2)) | (e2 = op(e0, e2))) & ((e2 = op(e2, e4)) | (e2 = op(e2, e3)) | (e2 = op(e2, e2)) | (e2 = op(e2, e1)) | (e2 = op(e2, e0))) & ((e1 = op(e4, e2)) | (e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(e0, e2))) & ((e1 = op(e2, e4)) | (e1 = op(e2, e3)) | (e1 = op(e2, e2)) | (e1 = op(e2, e1)) | (e1 = op(e2, e0))) & ((e0 = op(e4, e2)) | (e0 = op(e3, e2)) | (e0 = op(e2, e2)) | (e0 = op(e1, e2)) | (e0 = op(e0, e2))) & ((e0 = op(e2, e4)) | (e0 = op(e2, e3)) | (e0 = op(e2, e2)) | (e0 = op(e2, e1)) | (e0 = op(e2, e0))) & ((e4 = op(e4, e1)) | (e4 = op(e3, e1)) | (e4 = op(e2, e1)) | (e4 = op(e1, e1)) | (e4 = op(e0, e1))) & ((e4 = op(e1, e4)) | (e4 = op(e1, e3)) | (e4 = op(e1, e2)) | (e4 = op(e1, e1)) | (e4 = op(e1, e0))) & ((e3 = op(e4, e1)) | (e3 = op(e3, e1)) | (e3 = op(e2, e1)) | (e3 = op(e1, e1)) | (e3 = op(e0, e1))) & ((e3 = op(e1, e4)) | (e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, e0))) & ((e2 = op(e4, e1)) | (e2 = op(e3, e1)) | (e2 = op(e2, e1)) | (e2 = op(e1, e1)) | (e2 = op(e0, e1))) & ((e2 = op(e1, e4)) | (e2 = op(e1, e3)) | (e2 = op(e1, e2)) | (e2 = op(e1, e1)) | (e2 = op(e1, e0))) & ((e1 = op(e4, e1)) | (e1 = op(e3, e1)) | (e1 = op(e2, e1)) | (e1 = op(e1, e1)) | (e1 = op(e0, e1))) & ((e1 = op(e1, e4)) | (e1 = op(e1, e3)) | (e1 = op(e1, e2)) | (e1 = op(e1, e1)) | (e1 = op(e1, e0))) & ((e0 = op(e4, e1)) | (e0 = op(e3, e1)) | (e0 = op(e2, e1)) | (e0 = op(e1, e1)) | (e0 = op(e0, e1))) & ((e0 = op(e1, e4)) | (e0 = op(e1, e3)) | (e0 = op(e1, e2)) | (e0 = op(e1, e1)) | (e0 = op(e1, e0))) & ((e4 = op(e4, e0)) | (e4 = op(e3, e0)) | (e4 = op(e2, e0)) | (e4 = op(e1, e0)) | (e4 = op(e0, e0))) & ((e4 = op(e0, e4)) | (e4 = op(e0, e3)) | (e4 = op(e0, e2)) | (e4 = op(e0, e1)) | (e4 = op(e0, e0))) & ((e3 = op(e4, e0)) | (e3 = op(e3, e0)) | (e3 = op(e2, e0)) | (e3 = op(e1, e0)) | (e3 = op(e0, e0))) & ((e3 = op(e0, e4)) | (e3 = op(e0, e3)) | (e3 = op(e0, e2)) | (e3 = op(e0, e1)) | (e3 = op(e0, e0))) & ((e2 = op(e4, e0)) | (e2 = op(e3, e0)) | (e2 = op(e2, e0)) | (e2 = op(e1, e0)) | (e2 = op(e0, e0))) & ((e2 = op(e0, e4)) | (e2 = op(e0, e3)) | (e2 = op(e0, e2)) | (e2 = op(e0, e1)) | (e2 = op(e0, e0))) & ((e1 = op(e4, e0)) | (e1 = op(e3, e0)) | (e1 = op(e2, e0)) | (e1 = op(e1, e0)) | (e1 = op(e0, e0))) & ((e1 = op(e0, e4)) | (e1 = op(e0, e3)) | (e1 = op(e0, e2)) | (e1 = op(e0, e1)) | (e1 = op(e0, e0))) & ((e0 = op(e4, e0)) | (e0 = op(e3, e0)) | (e0 = op(e2, e0)) | (e0 = op(e1, e0)) | (e0 = op(e0, e0))) & ((e0 = op(e0, e4)) | (e0 = op(e0, e3)) | (e0 = op(e0, e2)) | (e0 = op(e0, e1)) | (e0 = op(e0, e0))) & ((e4 = unit) | (e3 = unit) | (e2 = unit) | (e1 = unit) | (e0 = unit)) & (e4 = op(e4, unit)) & (e4 = op(unit, e4)) & (e3 = op(e3, unit)) & (e3 = op(unit, e3)) & (e2 = op(e2, unit)) & (e2 = op(unit, e2)) & (e1 = op(e1, unit)) & (e1 = op(unit, e1)) & (e0 = op(e0, unit)) & (e0 = op(unit, e0)) & ((e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | (e1 = op(e4, e4)) | (e0 = op(e4, e4))) & ((e4 = op(e4, e3)) | (e3 = op(e4, e3)) | (e2 = op(e4, e3)) | (e1 = op(e4, e3)) | (e0 = op(e4, e3))) & ((e4 = op(e4, e2)) | (e3 = op(e4, e2)) | (e2 = op(e4, e2)) | (e1 = op(e4, e2)) | (e0 = op(e4, e2))) & ((e4 = op(e4, e1)) | (e3 = op(e4, e1)) | (e2 = op(e4, e1)) | (e1 = op(e4, e1)) | (e0 = op(e4, e1))) & ((e4 = op(e4, e0)) | (e3 = op(e4, e0)) | (e2 = op(e4, e0)) | (e1 = op(e4, e0)) | (e0 = op(e4, e0))) & ((e4 = op(e3, e4)) | (e3 = op(e3, e4)) | (e2 = op(e3, e4)) | (e1 = op(e3, e4)) | (e0 = op(e3, e4))) & ((e4 = op(e3, e3)) | (e3 = op(e3, e3)) | (e2 = op(e3, e3)) | (e1 = op(e3, e3)) | (e0 = op(e3, e3))) & ((e4 = op(e3, e2)) | (e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (e0 = op(e3, e2))) & ((e4 = op(e3, e1)) | (e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (e0 = op(e3, e1))) & ((e4 = op(e3, e0)) | (e3 = op(e3, e0)) | (e2 = op(e3, e0)) | (e1 = op(e3, e0)) | (e0 = op(e3, e0))) & ((e4 = op(e2, e4)) | (e3 = op(e2, e4)) | (e2 = op(e2, e4)) | (e1 = op(e2, e4)) | (e0 = op(e2, e4))) & ((e4 = op(e2, e3)) | (e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))) & ((e4 = op(e2, e2)) | (e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (e0 = op(e2, e2))) & ((e4 = op(e2, e1)) | (e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))) & ((e4 = op(e2, e0)) | (e3 = op(e2, e0)) | (e2 = op(e2, e0)) | (e1 = op(e2, e0)) | (e0 = op(e2, e0))) & ((e4 = op(e1, e4)) | (e3 = op(e1, e4)) | (e2 = op(e1, e4)) | (e1 = op(e1, e4)) | (e0 = op(e1, e4))) & ((e4 = op(e1, e3)) | (e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (e0 = op(e1, e3))) & ((e4 = op(e1, e2)) | (e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (e0 = op(e1, e2))) & ((e4 = op(e1, e1)) | (e3 = op(e1, e1)) | (e2 = op(e1, e1)) | (e1 = op(e1, e1)) | (e0 = op(e1, e1))) & ((e4 = op(e1, e0)) | (e3 = op(e1, e0)) | (e2 = op(e1, e0)) | (e1 = op(e1, e0)) | (e0 = op(e1, e0))) & ((e4 = op(e0, e4)) | (e3 = op(e0, e4)) | (e2 = op(e0, e4)) | (e1 = op(e0, e4)) | (e0 = op(e0, e4))) & ((e4 = op(e0, e3)) | (e3 = op(e0, e3)) | (e2 = op(e0, e3)) | (e1 = op(e0, e3)) | (e0 = op(e0, e3))) & ((e4 = op(e0, e2)) | (e3 = op(e0, e2)) | (e2 = op(e0, e2)) | (e1 = op(e0, e2)) | (e0 = op(e0, e2))) & ((e4 = op(e0, e1)) | (e3 = op(e0, e1)) | (e2 = op(e0, e1)) | (e1 = op(e0, e1)) | (e0 = op(e0, e1))) & ((e4 = op(e0, e0)) | (e3 = op(e0, e0)) | (e2 = op(e0, e0)) | (e1 = op(e0, e0)) | (e0 = op(e0, e0))) & (((e4 = op(e4, op(e4, e4))) & ~ (e4 = op(e4, op(e4, e4)))) | ((e3 = op(e4, op(e3, e4))) & ~ (e4 = op(e3, op(e3, e4)))) | ((e2 = op(e4, op(e2, e4))) & ~ (e4 = op(e2, op(e2, e4)))) | ((e1 = op(e4, op(e1, e4))) & ~ (e4 = op(e1, op(e1, e4)))) | ((e0 = op(e4, op(e0, e4))) & ~ (e4 = op(e0, op(e0, e4))))) & (((e4 = op(e3, op(e4, e3))) & ~ (e3 = op(e4, op(e4, e3)))) | ((e3 = op(e3, op(e3, e3))) & ~ (e3 = op(e3, op(e3, e3)))) | ((e2 = op(e3, op(e2, e3))) & ~ (e3 = op(e2, op(e2, e3)))) | ((e1 = op(e3, op(e1, e3))) & ~ (e3 = op(e1, op(e1, e3)))) | ((e0 = op(e3, op(e0, e3))) & ~ (e3 = op(e0, op(e0, e3))))) & (((e4 = op(e2, op(e4, e2))) & ~ (e2 = op(e4, op(e4, e2)))) | ((e3 = op(e2, op(e3, e2))) & ~ (e2 = op(e3, op(e3, e2)))) | ((e2 = op(e2, op(e2, e2))) & ~ (e2 = op(e2, op(e2, e2)))) | ((e1 = op(e2, op(e1, e2))) & ~ (e2 = op(e1, op(e1, e2)))) | ((e0 = op(e2, op(e0, e2))) & ~ (e2 = op(e0, op(e0, e2))))) & (((e4 = op(e1, op(e4, e1))) & ~ (e1 = op(e4, op(e4, e1)))) | ((e3 = op(e1, op(e3, e1))) & ~ (e1 = op(e3, op(e3, e1)))) | ((e2 = op(e1, op(e2, e1))) & ~ (e1 = op(e2, op(e2, e1)))) | ((e1 = op(e1, op(e1, e1))) & ~ (e1 = op(e1, op(e1, e1)))) | ((e0 = op(e1, op(e0, e1))) & ~ (e1 = op(e0, op(e0, e1))))) & (((e4 = op(e0, op(e4, e0))) & ~ (e0 = op(e4, op(e4, e0)))) | ((e3 = op(e0, op(e3, e0))) & ~ (e0 = op(e3, op(e3, e0)))) | ((e2 = op(e0, op(e2, e0))) & ~ (e0 = op(e2, op(e2, e0)))) | ((e1 = op(e0, op(e1, e0))) & ~ (e0 = op(e1, op(e1, e0)))) | ((e0 = op(e0, op(e0, e0))) & ~ (e0 = op(e0, op(e0, e0)))))), inference(negated_conjecture, [], [f4])).
fof(f4, plain, ~ (((e4 = op(e4, e4)) | (e4 = op(e3, e4)) | (e4 = op(e2, e4)) | (e4 = op(e1, e4)) | (e4 = op(e0, e4))) & ((e4 = op(e4, e4)) | (e4 = op(e4, e3)) | (e4 = op(e4, e2)) | (e4 = op(e4, e1)) | (e4 = op(e4, e0))) & ((e3 = op(e4, e4)) | (e3 = op(e3, e4)) | (e3 = op(e2, e4)) | (e3 = op(e1, e4)) | (e3 = op(e0, e4))) & ((e3 = op(e4, e4)) | (e3 = op(e4, e3)) | (e3 = op(e4, e2)) | (e3 = op(e4, e1)) | (e3 = op(e4, e0))) & ((e2 = op(e4, e4)) | (e2 = op(e3, e4)) | (e2 = op(e2, e4)) | (e2 = op(e1, e4)) | (e2 = op(e0, e4))) & ((e2 = op(e4, e4)) | (e2 = op(e4, e3)) | (e2 = op(e4, e2)) | (e2 = op(e4, e1)) | (e2 = op(e4, e0))) & ((e1 = op(e4, e4)) | (e1 = op(e3, e4)) | (e1 = op(e2, e4)) | (e1 = op(e1, e4)) | (e1 = op(e0, e4))) & ((e1 = op(e4, e4)) | (e1 = op(e4, e3)) | (e1 = op(e4, e2)) | (e1 = op(e4, e1)) | (e1 = op(e4, e0))) & ((e0 = op(e4, e4)) | (e0 = op(e3, e4)) | (e0 = op(e2, e4)) | (e0 = op(e1, e4)) | (e0 = op(e0, e4))) & ((e0 = op(e4, e4)) | (e0 = op(e4, e3)) | (e0 = op(e4, e2)) | (e0 = op(e4, e1)) | (e0 = op(e4, e0))) & ((e4 = op(e4, e3)) | (e4 = op(e3, e3)) | (e4 = op(e2, e3)) | (e4 = op(e1, e3)) | (e4 = op(e0, e3))) & ((e4 = op(e3, e4)) | (e4 = op(e3, e3)) | (e4 = op(e3, e2)) | (e4 = op(e3, e1)) | (e4 = op(e3, e0))) & ((e3 = op(e4, e3)) | (e3 = op(e3, e3)) | (e3 = op(e2, e3)) | (e3 = op(e1, e3)) | (e3 = op(e0, e3))) & ((e3 = op(e3, e4)) | (e3 = op(e3, e3)) | (e3 = op(e3, e2)) | (e3 = op(e3, e1)) | (e3 = op(e3, e0))) & ((e2 = op(e4, e3)) | (e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(e0, e3))) & ((e2 = op(e3, e4)) | (e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, e0))) & ((e1 = op(e4, e3)) | (e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(e0, e3))) & ((e1 = op(e3, e4)) | (e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, e0))) & ((e0 = op(e4, e3)) | (e0 = op(e3, e3)) | (e0 = op(e2, e3)) | (e0 = op(e1, e3)) | (e0 = op(e0, e3))) & ((e0 = op(e3, e4)) | (e0 = op(e3, e3)) | (e0 = op(e3, e2)) | (e0 = op(e3, e1)) | (e0 = op(e3, e0))) & ((e4 = op(e4, e2)) | (e4 = op(e3, e2)) | (e4 = op(e2, e2)) | (e4 = op(e1, e2)) | (e4 = op(e0, e2))) & ((e4 = op(e2, e4)) | (e4 = op(e2, e3)) | (e4 = op(e2, e2)) | (e4 = op(e2, e1)) | (e4 = op(e2, e0))) & ((e3 = op(e4, e2)) | (e3 = op(e3, e2)) | (e3 = op(e2, e2)) | (e3 = op(e1, e2)) | (e3 = op(e0, e2))) & ((e3 = op(e2, e4)) | (e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, e0))) & ((e2 = op(e4, e2)) | (e2 = op(e3, e2)) | (e2 = op(e2, e2)) | (e2 = op(e1, e2)) | (e2 = op(e0, e2))) & ((e2 = op(e2, e4)) | (e2 = op(e2, e3)) | (e2 = op(e2, e2)) | (e2 = op(e2, e1)) | (e2 = op(e2, e0))) & ((e1 = op(e4, e2)) | (e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(e0, e2))) & ((e1 = op(e2, e4)) | (e1 = op(e2, e3)) | (e1 = op(e2, e2)) | (e1 = op(e2, e1)) | (e1 = op(e2, e0))) & ((e0 = op(e4, e2)) | (e0 = op(e3, e2)) | (e0 = op(e2, e2)) | (e0 = op(e1, e2)) | (e0 = op(e0, e2))) & ((e0 = op(e2, e4)) | (e0 = op(e2, e3)) | (e0 = op(e2, e2)) | (e0 = op(e2, e1)) | (e0 = op(e2, e0))) & ((e4 = op(e4, e1)) | (e4 = op(e3, e1)) | (e4 = op(e2, e1)) | (e4 = op(e1, e1)) | (e4 = op(e0, e1))) & ((e4 = op(e1, e4)) | (e4 = op(e1, e3)) | (e4 = op(e1, e2)) | (e4 = op(e1, e1)) | (e4 = op(e1, e0))) & ((e3 = op(e4, e1)) | (e3 = op(e3, e1)) | (e3 = op(e2, e1)) | (e3 = op(e1, e1)) | (e3 = op(e0, e1))) & ((e3 = op(e1, e4)) | (e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, e0))) & ((e2 = op(e4, e1)) | (e2 = op(e3, e1)) | (e2 = op(e2, e1)) | (e2 = op(e1, e1)) | (e2 = op(e0, e1))) & ((e2 = op(e1, e4)) | (e2 = op(e1, e3)) | (e2 = op(e1, e2)) | (e2 = op(e1, e1)) | (e2 = op(e1, e0))) & ((e1 = op(e4, e1)) | (e1 = op(e3, e1)) | (e1 = op(e2, e1)) | (e1 = op(e1, e1)) | (e1 = op(e0, e1))) & ((e1 = op(e1, e4)) | (e1 = op(e1, e3)) | (e1 = op(e1, e2)) | (e1 = op(e1, e1)) | (e1 = op(e1, e0))) & ((e0 = op(e4, e1)) | (e0 = op(e3, e1)) | (e0 = op(e2, e1)) | (e0 = op(e1, e1)) | (e0 = op(e0, e1))) & ((e0 = op(e1, e4)) | (e0 = op(e1, e3)) | (e0 = op(e1, e2)) | (e0 = op(e1, e1)) | (e0 = op(e1, e0))) & ((e4 = op(e4, e0)) | (e4 = op(e3, e0)) | (e4 = op(e2, e0)) | (e4 = op(e1, e0)) | (e4 = op(e0, e0))) & ((e4 = op(e0, e4)) | (e4 = op(e0, e3)) | (e4 = op(e0, e2)) | (e4 = op(e0, e1)) | (e4 = op(e0, e0))) & ((e3 = op(e4, e0)) | (e3 = op(e3, e0)) | (e3 = op(e2, e0)) | (e3 = op(e1, e0)) | (e3 = op(e0, e0))) & ((e3 = op(e0, e4)) | (e3 = op(e0, e3)) | (e3 = op(e0, e2)) | (e3 = op(e0, e1)) | (e3 = op(e0, e0))) & ((e2 = op(e4, e0)) | (e2 = op(e3, e0)) | (e2 = op(e2, e0)) | (e2 = op(e1, e0)) | (e2 = op(e0, e0))) & ((e2 = op(e0, e4)) | (e2 = op(e0, e3)) | (e2 = op(e0, e2)) | (e2 = op(e0, e1)) | (e2 = op(e0, e0))) & ((e1 = op(e4, e0)) | (e1 = op(e3, e0)) | (e1 = op(e2, e0)) | (e1 = op(e1, e0)) | (e1 = op(e0, e0))) & ((e1 = op(e0, e4)) | (e1 = op(e0, e3)) | (e1 = op(e0, e2)) | (e1 = op(e0, e1)) | (e1 = op(e0, e0))) & ((e0 = op(e4, e0)) | (e0 = op(e3, e0)) | (e0 = op(e2, e0)) | (e0 = op(e1, e0)) | (e0 = op(e0, e0))) & ((e0 = op(e0, e4)) | (e0 = op(e0, e3)) | (e0 = op(e0, e2)) | (e0 = op(e0, e1)) | (e0 = op(e0, e0))) & ((e4 = unit) | (e3 = unit) | (e2 = unit) | (e1 = unit) | (e0 = unit)) & (e4 = op(e4, unit)) & (e4 = op(unit, e4)) & (e3 = op(e3, unit)) & (e3 = op(unit, e3)) & (e2 = op(e2, unit)) & (e2 = op(unit, e2)) & (e1 = op(e1, unit)) & (e1 = op(unit, e1)) & (e0 = op(e0, unit)) & (e0 = op(unit, e0)) & ((e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | (e1 = op(e4, e4)) | (e0 = op(e4, e4))) & ((e4 = op(e4, e3)) | (e3 = op(e4, e3)) | (e2 = op(e4, e3)) | (e1 = op(e4, e3)) | (e0 = op(e4, e3))) & ((e4 = op(e4, e2)) | (e3 = op(e4, e2)) | (e2 = op(e4, e2)) | (e1 = op(e4, e2)) | (e0 = op(e4, e2))) & ((e4 = op(e4, e1)) | (e3 = op(e4, e1)) | (e2 = op(e4, e1)) | (e1 = op(e4, e1)) | (e0 = op(e4, e1))) & ((e4 = op(e4, e0)) | (e3 = op(e4, e0)) | (e2 = op(e4, e0)) | (e1 = op(e4, e0)) | (e0 = op(e4, e0))) & ((e4 = op(e3, e4)) | (e3 = op(e3, e4)) | (e2 = op(e3, e4)) | (e1 = op(e3, e4)) | (e0 = op(e3, e4))) & ((e4 = op(e3, e3)) | (e3 = op(e3, e3)) | (e2 = op(e3, e3)) | (e1 = op(e3, e3)) | (e0 = op(e3, e3))) & ((e4 = op(e3, e2)) | (e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (e0 = op(e3, e2))) & ((e4 = op(e3, e1)) | (e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (e0 = op(e3, e1))) & ((e4 = op(e3, e0)) | (e3 = op(e3, e0)) | (e2 = op(e3, e0)) | (e1 = op(e3, e0)) | (e0 = op(e3, e0))) & ((e4 = op(e2, e4)) | (e3 = op(e2, e4)) | (e2 = op(e2, e4)) | (e1 = op(e2, e4)) | (e0 = op(e2, e4))) & ((e4 = op(e2, e3)) | (e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))) & ((e4 = op(e2, e2)) | (e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (e0 = op(e2, e2))) & ((e4 = op(e2, e1)) | (e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))) & ((e4 = op(e2, e0)) | (e3 = op(e2, e0)) | (e2 = op(e2, e0)) | (e1 = op(e2, e0)) | (e0 = op(e2, e0))) & ((e4 = op(e1, e4)) | (e3 = op(e1, e4)) | (e2 = op(e1, e4)) | (e1 = op(e1, e4)) | (e0 = op(e1, e4))) & ((e4 = op(e1, e3)) | (e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (e0 = op(e1, e3))) & ((e4 = op(e1, e2)) | (e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (e0 = op(e1, e2))) & ((e4 = op(e1, e1)) | (e3 = op(e1, e1)) | (e2 = op(e1, e1)) | (e1 = op(e1, e1)) | (e0 = op(e1, e1))) & ((e4 = op(e1, e0)) | (e3 = op(e1, e0)) | (e2 = op(e1, e0)) | (e1 = op(e1, e0)) | (e0 = op(e1, e0))) & ((e4 = op(e0, e4)) | (e3 = op(e0, e4)) | (e2 = op(e0, e4)) | (e1 = op(e0, e4)) | (e0 = op(e0, e4))) & ((e4 = op(e0, e3)) | (e3 = op(e0, e3)) | (e2 = op(e0, e3)) | (e1 = op(e0, e3)) | (e0 = op(e0, e3))) & ((e4 = op(e0, e2)) | (e3 = op(e0, e2)) | (e2 = op(e0, e2)) | (e1 = op(e0, e2)) | (e0 = op(e0, e2))) & ((e4 = op(e0, e1)) | (e3 = op(e0, e1)) | (e2 = op(e0, e1)) | (e1 = op(e0, e1)) | (e0 = op(e0, e1))) & ((e4 = op(e0, e0)) | (e3 = op(e0, e0)) | (e2 = op(e0, e0)) | (e1 = op(e0, e0)) | (e0 = op(e0, e0))) & (((e4 = op(e4, op(e4, e4))) & ~ (e4 = op(e4, op(e4, e4)))) | ((e3 = op(e4, op(e3, e4))) & ~ (e4 = op(e3, op(e3, e4)))) | ((e2 = op(e4, op(e2, e4))) & ~ (e4 = op(e2, op(e2, e4)))) | ((e1 = op(e4, op(e1, e4))) & ~ (e4 = op(e1, op(e1, e4)))) | ((e0 = op(e4, op(e0, e4))) & ~ (e4 = op(e0, op(e0, e4))))) & (((e4 = op(e3, op(e4, e3))) & ~ (e3 = op(e4, op(e4, e3)))) | ((e3 = op(e3, op(e3, e3))) & ~ (e3 = op(e3, op(e3, e3)))) | ((e2 = op(e3, op(e2, e3))) & ~ (e3 = op(e2, op(e2, e3)))) | ((e1 = op(e3, op(e1, e3))) & ~ (e3 = op(e1, op(e1, e3)))) | ((e0 = op(e3, op(e0, e3))) & ~ (e3 = op(e0, op(e0, e3))))) & (((e4 = op(e2, op(e4, e2))) & ~ (e2 = op(e4, op(e4, e2)))) | ((e3 = op(e2, op(e3, e2))) & ~ (e2 = op(e3, op(e3, e2)))) | ((e2 = op(e2, op(e2, e2))) & ~ (e2 = op(e2, op(e2, e2)))) | ((e1 = op(e2, op(e1, e2))) & ~ (e2 = op(e1, op(e1, e2)))) | ((e0 = op(e2, op(e0, e2))) & ~ (e2 = op(e0, op(e0, e2))))) & (((e4 = op(e1, op(e4, e1))) & ~ (e1 = op(e4, op(e4, e1)))) | ((e3 = op(e1, op(e3, e1))) & ~ (e1 = op(e3, op(e3, e1)))) | ((e2 = op(e1, op(e2, e1))) & ~ (e1 = op(e2, op(e2, e1)))) | ((e1 = op(e1, op(e1, e1))) & ~ (e1 = op(e1, op(e1, e1)))) | ((e0 = op(e1, op(e0, e1))) & ~ (e1 = op(e0, op(e0, e1))))) & (((e4 = op(e0, op(e4, e0))) & ~ (e0 = op(e4, op(e4, e0)))) | ((e3 = op(e0, op(e3, e0))) & ~ (e0 = op(e3, op(e3, e0)))) | ((e2 = op(e0, op(e2, e0))) & ~ (e0 = op(e2, op(e2, e0)))) | ((e1 = op(e0, op(e1, e0))) & ~ (e0 = op(e1, op(e1, e0)))) | ((e0 = op(e0, op(e0, e0))) & ~ (e0 = op(e0, op(e0, e0)))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG092+1.p', co1)).