fof(f10921, plain, $false, inference(avatar_sat_refutation, [], [f1220, f1241, f1262, f1304, f1325, f1346, f1367, f1388, f1409, f1430, f1451, f1472, f1493, f1514, f1535, f1556, f1577, f1598, f1619, f1640, f1661, f1682, f1703, f1724, f1745, f1748, f1749, f1750, f1751, f1752, f1753, f1754, f1755, f1756, f1757, f1760, f1761, f1762, f1763, f1764, f1765, f1766, f1767, f1768, f1769, f1772, f1773, f1774, f1775, f1776, f1777, f1778, f1779, f1780, f1781, f1784, f1785, f1786, f1787, f1788, f1789, f1790, f1791, f1792, f1793, f1804, f1809, f1814, f1819, f1824, f1834, f1839, f1844, f1859, f1864, f1869, f1874, f1884, f1889, f1894, f1899, f1909, f1914, f1919, f1924, f1933, f1943, f1953, f1958, f1963, f1968, f1973, f1978, f1983, f1993, f1998, f2003, f2008, f2018, f2023, f2028, f2033, f2043, f2048, f2053, f2062, f2067, f2082, f2087, f2092, f2107, f2112, f2117, f2122, f2127, f2132, f2137, f2142, f2152, f2157, f2162, f2167, f2177, f2182, f2191, f2196, f2216, f2221, f2226, f2236, f2241, f2246, f2251, f2261, f2266, f2271, f2276, f2281, f2286, f2291, f2296, f2301, f2311, f2320, f2325, f2330, f2335, f2345, f2350, f2355, f2360, f2370, f2375, f2380, f2385, f2395, f2400, f2405, f2410, f2420, f2425, f2430, f2435, f2440, f2441, f2443, f2445, f2446, f2447, f2449, f2452, f2454, f2456, f2461, f2463, f2465, f2467, f2470, f2472, f2474, f2476, f2479, f2481, f2483, f2485, f2487, f2490, f2494, f2496, f2497, f2499, f2501, f2503, f2505, f2508, f2510, f2512, f2514, f2517, f2519, f2520, f2521, f2523, f2526, f2528, f2529, f2530, f2532, f2534, f2539, f2541, f2543, f2548, f2550, f2552, f2553, f2555, f2557, f2559, f2561, f2564, f2566, f2568, f2570, f2573, f2574, f2575, f2577, f2579, f2586, f2588, f2590, f2593, f2595, f2597, f2599, f2602, f2604, f2606, f2608, f2609, f2611, f2613, f2615, f2617, f2620, f2622, f2624, f2626, f2628, f2631, f2633, f2635, f2637, f2640, f2642, f2644, f2646, f2649, f2651, f2653, f2655, f2658, f2660, f2662, f2664, f2665, f2666, f2684, f2702, f2733, f2738, f2768, f2778, f2789, f2790, f2831, f2832, f2845, f2850, f2873, f2879, f2885, f2887, f2900, f2911, f2921, f2922, f2927, f2932, f2942, f2945, f2958, f2964, f2969, f2974, f2975, f2976, f2981, f2982, f2998, f3004, f3009, f3010, f3016, f3019, f3021, f3026, f3069, f3070, f3083, f3084, f3120, f3123, f3126, f3135, f3136, f3155, f3170, f3192, f3195, f3206, f3210, f3246, f3259, f3267, f3277, f3307, f3316, f3317, f3322, f3325, f3330, f3334, f3337, f3344, f3347, f3350, f3353, f3361, f3367, f3375, f3383, f3394, f3405, f3408, f3410, f3429, f3465, f3468, f3469, f3470, f3484, f3517, f3520, f3527, f3533, f3557, f3564, f3565, f3575, f3584, f3625, f3651, f3655, f3657, f3674, f3690, f3696, f3753, f3762, f3763, f3766, f3780, f3787, f3791, f3803, f3807, f3827, f3831, f3875, f3942, f3947, f3962, f3994, f4008, f4012, f4022, f4040, f4043, f4051, f4080, f4088, f4125, f4131, f4155, f4175, f4195, f4235, f4242, f4261, f4262, f4268, f4272, f4305, f4329, f4349, f4358, f4385, f4399, f4401, f4413, f4424, f4437, f4494, f4496, f4516, f4605, f4607, f4612, f4616, f4622, f4643, f4646, f4685, f4725, f4732, f4765, f4788, f4800, f4825, f4829, f4860, f4862, f4866, f4872, f4914, f4946, f5049, f5068, f5090, f5093, f5136, f5168, f5224, f5245, f5269, f5273, f5278, f5285, f5308, f5369, f5370, f5372, f5381, f5393, f5399, f5446, f5459, f5467, f5469, f5479, f5502, f5565, f5573, f5579, f5589, f5593, f5601, f5610, f5611, f5667, f5678, f5681, f5744, f5776, f5807, f5826, f5848, f5866, f5880, f5887, f5907, f5908, f5909, f5927, f5939, f6006, f6012, f6031, f6041, f6081, f6089, f6107, f6164, f6172, f6193, f6197, f6208, f6209, f6276, f6282, f6293, f6342, f6343, f6345, f6347, f6382, f6391, f6441, f6506, f6529, f6553, f6591, f6610, f6623, f6633, f6634, f6637, f6655, f6664, f6669, f6672, f6697, f6701, f6721, f6728, f6737, f6755, f6761, f6777, f6786, f6800, f6852, f6856, f6870, f6905, f6916, f6947, f6956, f6962, f6975, f7055, f7066, f7073, f7089, f7098, f7099, f7103, f7114, f7115, f7146, f7188, f7190, f7192, f7235, f7255, f7267, f7293, f7302, f7312, f7321, f7342, f7366, f7393, f7395, f7400, f7413, f7439, f7441, f7475, f7483, f7484, f7501, f7522, f7541, f7550, f7551, f7565, f7595, f7646, f7662, f7700, f7716, f7723, f7737, f7750, f7758, f7766, f7767, f7781, f7784, f7785, f7799, f7808, f7840, f7872, f7873, f7874, f7875, f7876, f7877, f7879, f7880, f7881, f7896, f7897, f7899, f7916, f7917, f7931, f7939, f7945, f7955, f7967, f7979, f7980, f7987, f7991, f8001, f8013, f8022, f8035, f8048, f8054, f8068, f8095, f8098, f8103, f8111, f8135, f8136, f8137, f8138, f8139, f8140, f8141, f8142, f8157, f8163, f8172, f8190, f8204, f8214, f8233, f8242, f8249, f8257, f8297, f8304, f8343, f8344, f8345, f8346, f8354, f8366, f8374, f8391, f8409, f8419, f8432, f8443, f8453, f8462, f8463, f8470, f8479, f8483, f8488, f8510, f8514, f8554, f8588, f8590, f8628, f8629, f8632, f8639, f8648, f8660, f8678, f8683, f8731, f8732, f8778, f8784, f8794, f8797, f8804, f8805, f8873, f8889, f8891, f8894, f8916, f8941, f8942, f8944, f8950, f8962, f8972, f8984, f9005, f9020, f9026, f9028, f9045, f9155, f9161, f9162, f9178, f9264, f9266, f9268, f9313, f9319, f9336, f9337, f9346, f9350, f9358, f9366, f9375, f9408, f9411, f9413, f9415, f9449, f9450, f9453, f9461, f9495, f9501, f9556, f9561, f9591, f9599, f9600, f9652, f9654, f9681, f9688, f9702, f9745, f9746, f9767, f9785, f9811, f9819, f9820, f9863, f9864, f9866, f9868, f9876, f9887, f9888, f9893, f9918, f9936, f9943, f9948, f9993, f10014, f10022, f10028, f10034, f10046, f10067, f10068, f10076, f10082, f10091, f10097, f10098, f10112, f10124, f10134, f10141, f10147, f10163, f10169, f10180, f10199, f10208, f10226, f10227, f10228, f10229, f10231, f10232, f10233, f10234, f10235, f10261, f10262, f10281, f10298, f10299, f10318, f10330, f10390, f10432, f10461, f10462, f10463, f10464, f10465, f10466, f10467, f10468, f10469, f10488, f10507, f10508, f10517, f10525, f10530, f10542, f10557, f10558, f10560, f10566, f10567, f10576, f10582, f10591, f10619, f10629, f10641, f10660, f10685, f10716, f10730, f10731, f10733, f10734, f10735, f10736, f10737, f10738, f10739, f10756, f10757, f10761, f10762, f10774, f10780, f10781, f10788, f10794, f10806, f10854, f10877, f10878, f10887, f10907])).
fof(f10907, plain, (~ spl130_1 | ~ spl130_6), inference(avatar_split_clause, [], [f10902, f1222, f1201])).
fof(f1201, plain, (spl130_1 <=> (e0 = op(e4, e4))), introduced(avatar_definition, [new_symbols(naming, [spl130_1])])).
fof(f1222, plain, (spl130_6 <=> (e0 = op(e4, e3))), introduced(avatar_definition, [new_symbols(naming, [spl130_6])])).
fof(f10902, plain, (~ (e0 = op(e4, e4)) | ~ spl130_6), inference(backward_demodulation, [], [f693, f1224])).
fof(f1224, plain, ((e0 = op(e4, e3)) | ~ spl130_6), inference(avatar_component_clause, [], [f1222])).
fof(f693, plain, ~ (op(e4, e3) = op(e4, e4)), inference(cnf_transformation, [], [f4])).
fof(f4, plain, (~ (op(e4, e3) = op(e4, e4)) & ~ (op(e4, e2) = op(e4, e4)) & ~ (op(e4, e2) = op(e4, e3)) & ~ (op(e4, e1) = op(e4, e4)) & ~ (op(e4, e1) = op(e4, e3)) & ~ (op(e4, e1) = op(e4, e2)) & ~ (op(e4, e0) = op(e4, e4)) & ~ (op(e4, e0) = op(e4, e3)) & ~ (op(e4, e0) = op(e4, e2)) & ~ (op(e4, e0) = op(e4, e1)) & ~ (op(e3, e3) = op(e3, e4)) & ~ (op(e3, e2) = op(e3, e4)) & ~ (op(e3, e2) = op(e3, e3)) & ~ (op(e3, e1) = op(e3, e4)) & ~ (op(e3, e1) = op(e3, e3)) & ~ (op(e3, e1) = op(e3, e2)) & ~ (op(e3, e0) = op(e3, e4)) & ~ (op(e3, e0) = op(e3, e3)) & ~ (op(e3, e0) = op(e3, e2)) & ~ (op(e3, e0) = op(e3, e1)) & ~ (op(e2, e3) = op(e2, e4)) & ~ (op(e2, e2) = op(e2, e4)) & ~ (op(e2, e2) = op(e2, e3)) & ~ (op(e2, e1) = op(e2, e4)) & ~ (op(e2, e1) = op(e2, e3)) & ~ (op(e2, e1) = op(e2, e2)) & ~ (op(e2, e0) = op(e2, e4)) & ~ (op(e2, e0) = op(e2, e3)) & ~ (op(e2, e0) = op(e2, e2)) & ~ (op(e2, e0) = op(e2, e1)) & ~ (op(e1, e3) = op(e1, e4)) & ~ (op(e1, e2) = op(e1, e4)) & ~ (op(e1, e2) = op(e1, e3)) & ~ (op(e1, e1) = op(e1, e4)) & ~ (op(e1, e1) = op(e1, e3)) & ~ (op(e1, e1) = op(e1, e2)) & ~ (op(e1, e0) = op(e1, e4)) & ~ (op(e1, e0) = op(e1, e3)) & ~ (op(e1, e0) = op(e1, e2)) & ~ (op(e1, e0) = op(e1, e1)) & ~ (op(e0, e3) = op(e0, e4)) & ~ (op(e0, e2) = op(e0, e4)) & ~ (op(e0, e2) = op(e0, e3)) & ~ (op(e0, e1) = op(e0, e4)) & ~ (op(e0, e1) = op(e0, e3)) & ~ (op(e0, e1) = op(e0, e2)) & ~ (op(e0, e0) = op(e0, e4)) & ~ (op(e0, e0) = op(e0, e3)) & ~ (op(e0, e0) = op(e0, e2)) & ~ (op(e0, e0) = op(e0, e1)) & ~ (op(e3, e4) = op(e4, e4)) & ~ (op(e2, e4) = op(e4, e4)) & ~ (op(e2, e4) = op(e3, e4)) & ~ (op(e1, e4) = op(e4, e4)) & ~ (op(e1, e4) = op(e3, e4)) & ~ (op(e1, e4) = op(e2, e4)) & ~ (op(e0, e4) = op(e4, e4)) & ~ (op(e0, e4) = op(e3, e4)) & ~ (op(e0, e4) = op(e2, e4)) & ~ (op(e0, e4) = op(e1, e4)) & ~ (op(e3, e3) = op(e4, e3)) & ~ (op(e2, e3) = op(e4, e3)) & ~ (op(e2, e3) = op(e3, e3)) & ~ (op(e1, e3) = op(e4, e3)) & ~ (op(e1, e3) = op(e3, e3)) & ~ (op(e1, e3) = op(e2, e3)) & ~ (op(e0, e3) = op(e4, e3)) & ~ (op(e0, e3) = op(e3, e3)) & ~ (op(e0, e3) = op(e2, e3)) & ~ (op(e0, e3) = op(e1, e3)) & ~ (op(e3, e2) = op(e4, e2)) & ~ (op(e2, e2) = op(e4, e2)) & ~ (op(e2, e2) = op(e3, e2)) & ~ (op(e1, e2) = op(e4, e2)) & ~ (op(e1, e2) = op(e3, e2)) & ~ (op(e1, e2) = op(e2, e2)) & ~ (op(e0, e2) = op(e4, e2)) & ~ (op(e0, e2) = op(e3, e2)) & ~ (op(e0, e2) = op(e2, e2)) & ~ (op(e0, e2) = op(e1, e2)) & ~ (op(e3, e1) = op(e4, e1)) & ~ (op(e2, e1) = op(e4, e1)) & ~ (op(e2, e1) = op(e3, e1)) & ~ (op(e1, e1) = op(e4, e1)) & ~ (op(e1, e1) = op(e3, e1)) & ~ (op(e1, e1) = op(e2, e1)) & ~ (op(e0, e1) = op(e4, e1)) & ~ (op(e0, e1) = op(e3, e1)) & ~ (op(e0, e1) = op(e2, e1)) & ~ (op(e0, e1) = op(e1, e1)) & ~ (op(e3, e0) = op(e4, e0)) & ~ (op(e2, e0) = op(e4, e0)) & ~ (op(e2, e0) = op(e3, e0)) & ~ (op(e1, e0) = op(e4, e0)) & ~ (op(e1, e0) = op(e3, e0)) & ~ (op(e1, e0) = op(e2, e0)) & ~ (op(e0, e0) = op(e4, e0)) & ~ (op(e0, e0) = op(e3, e0)) & ~ (op(e0, e0) = op(e2, e0)) & ~ (op(e0, e0) = op(e1, e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG099+1.p', ax4)).
fof(f10887, plain, (~ spl130_27 | ~ spl130_37), inference(avatar_split_clause, [], [f10883, f1352, f1310])).
fof(f1310, plain, (spl130_27 <=> (e1 = op(e3, e4))), introduced(avatar_definition, [new_symbols(naming, [spl130_27])])).
fof(f1352, plain, (spl130_37 <=> (e1 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl130_37])])).
fof(f10883, plain, (~ (e1 = op(e3, e4)) | ~ spl130_37), inference(backward_demodulation, [], [f682, f1354])).
fof(f1354, plain, ((e1 = op(e3, e2)) | ~ spl130_37), inference(avatar_component_clause, [], [f1352])).
fof(f682, plain, ~ (op(e3, e2) = op(e3, e4)), inference(cnf_transformation, [], [f4])).
fof(f10878, plain, (~ spl130_45 | ~ spl130_50), inference(avatar_split_clause, [], [f10874, f1406, f1385])).
fof(f1385, plain, (spl130_45 <=> (e4 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl130_45])])).
fof(f1406, plain, (spl130_50 <=> (e4 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl130_50])])).
fof(f10874, plain, (~ (e4 = op(e3, e1)) | ~ spl130_50), inference(backward_demodulation, [], [f674, f1408])).
fof(f1408, plain, ((e4 = op(e3, e0)) | ~ spl130_50), inference(avatar_component_clause, [], [f1406])).
fof(f674, plain, ~ (op(e3, e0) = op(e3, e1)), inference(cnf_transformation, [], [f4])).
fof(f10877, plain, (~ spl130_25 | ~ spl130_50), inference(avatar_split_clause, [], [f10873, f1406, f1301])).
fof(f1301, plain, (spl130_25 <=> (e4 = op(e4, e0))), introduced(avatar_definition, [new_symbols(naming, [spl130_25])])).
fof(f10873, plain, (~ (e4 = op(e4, e0)) | ~ spl130_50), inference(backward_demodulation, [], [f603, f1408])).
fof(f603, plain, ~ (op(e3, e0) = op(e4, e0)), inference(cnf_transformation, [], [f4])).
fof(f10854, plain, (~ spl130_62 | ~ spl130_72), inference(avatar_split_clause, [], [f10850, f1499, f1457])).
fof(f1457, plain, (spl130_62 <=> (e1 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl130_62])])).
fof(f1499, plain, (spl130_72 <=> (e1 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl130_72])])).
fof(f10850, plain, (~ (e1 = op(e2, e2)) | ~ spl130_72), inference(backward_demodulation, [], [f665, f1501])).
fof(f1501, plain, ((e1 = op(e2, e0)) | ~ spl130_72), inference(avatar_component_clause, [], [f1499])).
fof(f665, plain, ~ (op(e2, e0) = op(e2, e2)), inference(cnf_transformation, [], [f4])).
fof(f10806, plain, (~ spl130_81 | ~ spl130_96), inference(avatar_split_clause, [], [f10800, f1600, f1537])).
fof(f1537, plain, (spl130_81 <=> (e0 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl130_81])])).
fof(f1600, plain, (spl130_96 <=> (e0 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl130_96])])).
fof(f10800, plain, (~ (e0 = op(e1, e3)) | ~ spl130_96), inference(backward_demodulation, [], [f656, f1602])).
fof(f1602, plain, ((e0 = op(e1, e0)) | ~ spl130_96), inference(avatar_component_clause, [], [f1600])).
fof(f656, plain, ~ (op(e1, e0) = op(e1, e3)), inference(cnf_transformation, [], [f4])).
fof(f10794, plain, (~ spl130_78 | ~ spl130_103), inference(avatar_split_clause, [], [f10789, f1629, f1524])).
fof(f1524, plain, (spl130_78 <=> (e2 = op(e1, e4))), introduced(avatar_definition, [new_symbols(naming, [spl130_78])])).
fof(f1629, plain, (spl130_103 <=> (e2 = op(e0, e4))), introduced(avatar_definition, [new_symbols(naming, [spl130_103])])).
fof(f10789, plain, (~ (e2 = op(e1, e4)) | ~ spl130_103), inference(backward_demodulation, [], [f634, f1631])).
fof(f1631, plain, ((e2 = op(e0, e4)) | ~ spl130_103), inference(avatar_component_clause, [], [f1629])).
fof(f634, plain, ~ (op(e0, e4) = op(e1, e4)), inference(cnf_transformation, [], [f4])).
fof(f10788, plain, (~ spl130_7 | ~ spl130_107), inference(avatar_split_clause, [], [f10783, f1646, f1226])).
fof(f1226, plain, (spl130_7 <=> (e1 = op(e4, e3))), introduced(avatar_definition, [new_symbols(naming, [spl130_7])])).
fof(f1646, plain, (spl130_107 <=> (e1 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl130_107])])).
fof(f10783, plain, (~ (e1 = op(e4, e3)) | ~ spl130_107), inference(backward_demodulation, [], [f627, f1648])).
fof(f1648, plain, ((e1 = op(e0, e3)) | ~ spl130_107), inference(avatar_component_clause, [], [f1646])).
fof(f627, plain, ~ (op(e0, e3) = op(e4, e3)), inference(cnf_transformation, [], [f4])).
fof(f10781, plain, (~ spl130_105 | ~ spl130_115), inference(avatar_split_clause, [], [f10779, f1679, f1637])).
fof(f1637, plain, (spl130_105 <=> (e4 = op(e0, e4))), introduced(avatar_definition, [new_symbols(naming, [spl130_105])])).
fof(f1679, plain, (spl130_115 <=> (e4 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl130_115])])).
fof(f10779, plain, (~ (e4 = op(e0, e4)) | ~ spl130_115), inference(backward_demodulation, [], [f652, f1681])).
fof(f1681, plain, ((e4 = op(e0, e2)) | ~ spl130_115), inference(avatar_component_clause, [], [f1679])).
fof(f652, plain, ~ (op(e0, e2) = op(e0, e4)), inference(cnf_transformation, [], [f4])).
fof(f10780, plain, (~ spl130_90 | ~ spl130_115), inference(avatar_split_clause, [], [f10775, f1679, f1574])).
fof(f1574, plain, (spl130_90 <=> (e4 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl130_90])])).
fof(f10775, plain, (~ (e4 = op(e1, e2)) | ~ spl130_115), inference(backward_demodulation, [], [f614, f1681])).
fof(f614, plain, ~ (op(e0, e2) = op(e1, e2)), inference(cnf_transformation, [], [f4])).
fof(f10774, plain, (~ spl130_66 | ~ spl130_116), inference(avatar_split_clause, [], [f10764, f1684, f1474])).
fof(f1474, plain, (spl130_66 <=> (e0 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl130_66])])).
fof(f1684, plain, (spl130_116 <=> (e0 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl130_116])])).
fof(f10764, plain, (~ (e0 = op(e2, e1)) | ~ spl130_116), inference(backward_demodulation, [], [f605, f1686])).
fof(f1686, plain, ((e0 = op(e0, e1)) | ~ spl130_116), inference(avatar_component_clause, [], [f1684])).
fof(f605, plain, ~ (op(e0, e1) = op(e2, e1)), inference(cnf_transformation, [], [f4])).
fof(f10762, plain, (spl130_50 | ~ spl130_124 | ~ spl130_304), inference(avatar_split_clause, [], [f10753, f2882, f1717, f1406])).
fof(f1717, plain, (spl130_124 <=> (op(e0, e0) = e3)), introduced(avatar_definition, [new_symbols(naming, [spl130_124])])).
fof(f2882, plain, (spl130_304 <=> (e4 = op(op(e0, e0), e0))), introduced(avatar_definition, [new_symbols(naming, [spl130_304])])).
fof(f10753, plain, ((e4 = op(e3, e0)) | (~ spl130_124 | ~ spl130_304)), inference(backward_demodulation, [], [f2883, f1719])).
fof(f1719, plain, ((op(e0, e0) = e3) | ~ spl130_124), inference(avatar_component_clause, [], [f1717])).
fof(f2883, plain, ((e4 = op(op(e0, e0), e0)) | ~ spl130_304), inference(avatar_component_clause, [], [f2882])).
fof(f10761, plain, (~ spl130_33 | ~ spl130_124 | spl130_302), inference(avatar_contradiction_clause, [], [f10760])).
fof(f10760, plain, ($false | (~ spl130_33 | ~ spl130_124 | spl130_302)), inference(subsumption_resolution, [], [f10752, f1337])).
fof(f1337, plain, ((e2 = op(e3, e3)) | ~ spl130_33), inference(avatar_component_clause, [], [f1335])).
fof(f1335, plain, (spl130_33 <=> (e2 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl130_33])])).
fof(f10752, plain, (~ (e2 = op(e3, e3)) | (~ spl130_124 | spl130_302)), inference(backward_demodulation, [], [f2871, f1719])).
fof(f2871, plain, (~ (e2 = op(op(e0, e0), op(e0, e0))) | spl130_302), inference(avatar_component_clause, [], [f2869])).
fof(f2869, plain, (spl130_302 <=> (e2 = op(op(e0, e0), op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl130_302])])).
fof(f10757, plain, (~ spl130_109 | ~ spl130_124), inference(avatar_split_clause, [], [f10746, f1717, f1654])).
fof(f1654, plain, (spl130_109 <=> (e3 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl130_109])])).
fof(f10746, plain, (~ (e3 = op(e0, e3)) | ~ spl130_124), inference(backward_demodulation, [], [f646, f1719])).
fof(f646, plain, ~ (op(e0, e0) = op(e0, e3)), inference(cnf_transformation, [], [f4])).
fof(f10756, plain, (~ spl130_49 | ~ spl130_124), inference(avatar_split_clause, [], [f10742, f1717, f1402])).
fof(f1402, plain, (spl130_49 <=> (e3 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl130_49])])).
fof(f10742, plain, (~ (e3 = op(e3, e0)) | ~ spl130_124), inference(backward_demodulation, [], [f596, f1719])).
fof(f596, plain, ~ (op(e0, e0) = op(e3, e0)), inference(cnf_transformation, [], [f4])).
fof(f10739, plain, (spl130_20 | ~ spl130_127), inference(avatar_split_clause, [], [f10726, f1730, f1280])).
fof(f1280, plain, (spl130_20 <=> (e4 = op(e4, e1))), introduced(avatar_definition, [new_symbols(naming, [spl130_20])])).
fof(f1730, plain, (spl130_127 <=> (e1 = unit)), introduced(avatar_definition, [new_symbols(naming, [spl130_127])])).
fof(f10726, plain, ((e4 = op(e4, e1)) | ~ spl130_127), inference(backward_demodulation, [], [f542, f1732])).
fof(f1732, plain, ((e1 = unit) | ~ spl130_127), inference(avatar_component_clause, [], [f1730])).
fof(f542, plain, (e4 = op(e4, unit)), inference(cnf_transformation, [], [f2])).
fof(f2, plain, (((e4 = unit) | (e3 = unit) | (e2 = unit) | (e1 = unit) | (e0 = unit)) & (e4 = op(e4, unit)) & (e4 = op(unit, e4)) & (e3 = op(e3, unit)) & (e3 = op(unit, e3)) & (e2 = op(e2, unit)) & (e2 = op(unit, e2)) & (e1 = op(e1, unit)) & (e1 = op(unit, e1)) & (e0 = op(e0, unit)) & (e0 = op(unit, e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG099+1.p', ax2)).
fof(f10738, plain, (spl130_80 | ~ spl130_127), inference(avatar_split_clause, [], [f10725, f1730, f1532])).
fof(f1532, plain, (spl130_80 <=> (e4 = op(e1, e4))), introduced(avatar_definition, [new_symbols(naming, [spl130_80])])).
fof(f10725, plain, ((e4 = op(e1, e4)) | ~ spl130_127), inference(backward_demodulation, [], [f541, f1732])).
fof(f541, plain, (e4 = op(unit, e4)), inference(cnf_transformation, [], [f2])).
fof(f10737, plain, (spl130_44 | ~ spl130_127), inference(avatar_split_clause, [], [f10724, f1730, f1381])).
fof(f1381, plain, (spl130_44 <=> (e3 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl130_44])])).
fof(f10724, plain, ((e3 = op(e3, e1)) | ~ spl130_127), inference(backward_demodulation, [], [f540, f1732])).
fof(f540, plain, (e3 = op(e3, unit)), inference(cnf_transformation, [], [f2])).
fof(f10736, plain, (spl130_84 | ~ spl130_127), inference(avatar_split_clause, [], [f10723, f1730, f1549])).
fof(f1549, plain, (spl130_84 <=> (e3 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl130_84])])).
fof(f10723, plain, ((e3 = op(e1, e3)) | ~ spl130_127), inference(backward_demodulation, [], [f539, f1732])).
fof(f539, plain, (e3 = op(unit, e3)), inference(cnf_transformation, [], [f2])).
fof(f10735, plain, (spl130_68 | ~ spl130_127), inference(avatar_split_clause, [], [f10722, f1730, f1482])).
fof(f1482, plain, (spl130_68 <=> (e2 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl130_68])])).
fof(f10722, plain, ((e2 = op(e2, e1)) | ~ spl130_127), inference(backward_demodulation, [], [f538, f1732])).
fof(f538, plain, (e2 = op(e2, unit)), inference(cnf_transformation, [], [f2])).
fof(f10734, plain, (spl130_88 | ~ spl130_127), inference(avatar_split_clause, [], [f10721, f1730, f1566])).
fof(f1566, plain, (spl130_88 <=> (e2 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl130_88])])).
fof(f10721, plain, ((e2 = op(e1, e2)) | ~ spl130_127), inference(backward_demodulation, [], [f537, f1732])).
fof(f537, plain, (e2 = op(unit, e2)), inference(cnf_transformation, [], [f2])).
fof(f10733, plain, (spl130_92 | ~ spl130_127), inference(avatar_split_clause, [], [f10720, f1730, f1583])).
fof(f1583, plain, (spl130_92 <=> (e1 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl130_92])])).
fof(f10720, plain, ((e1 = op(e1, e1)) | ~ spl130_127), inference(backward_demodulation, [], [f536, f1732])).
fof(f536, plain, (e1 = op(e1, unit)), inference(cnf_transformation, [], [f2])).
fof(f10731, plain, (spl130_116 | ~ spl130_127), inference(avatar_split_clause, [], [f10718, f1730, f1684])).
fof(f10718, plain, ((e0 = op(e0, e1)) | ~ spl130_127), inference(backward_demodulation, [], [f534, f1732])).
fof(f534, plain, (e0 = op(e0, unit)), inference(cnf_transformation, [], [f2])).
fof(f10730, plain, (spl130_96 | ~ spl130_127), inference(avatar_split_clause, [], [f10717, f1730, f1600])).
fof(f10717, plain, ((e0 = op(e1, e0)) | ~ spl130_127), inference(backward_demodulation, [], [f533, f1732])).
fof(f533, plain, (e0 = op(unit, e0)), inference(cnf_transformation, [], [f2])).
fof(f10716, plain, (~ spl130_62 | ~ spl130_33 | spl130_278), inference(avatar_split_clause, [], [f10715, f2745, f1335, f1457])).
fof(f2745, plain, (spl130_278 <=> (e1 = op(op(e3, e3), op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl130_278])])).
fof(f10715, plain, (~ (e1 = op(e2, e2)) | (~ spl130_33 | spl130_278)), inference(forward_demodulation, [], [f2747, f1337])).
fof(f2747, plain, (~ (e1 = op(op(e3, e3), op(e3, e3))) | spl130_278), inference(avatar_component_clause, [], [f2745])).
fof(f10685, plain, (~ spl130_2 | ~ spl130_287 | spl130_289), inference(avatar_split_clause, [], [f10437, f2804, f2792, f1205])).
fof(f1205, plain, (spl130_2 <=> (e1 = op(e4, e4))), introduced(avatar_definition, [new_symbols(naming, [spl130_2])])).
fof(f2792, plain, (spl130_287 <=> (e4 = op(op(e2, e2), e2))), introduced(avatar_definition, [new_symbols(naming, [spl130_287])])).
fof(f2804, plain, (spl130_289 <=> (e1 = op(op(op(e2, e2), e2), op(op(e2, e2), e2)))), introduced(avatar_definition, [new_symbols(naming, [spl130_289])])).
fof(f10437, plain, (~ (e1 = op(e4, e4)) | (~ spl130_287 | spl130_289)), inference(backward_demodulation, [], [f2806, f2793])).
fof(f2793, plain, ((e4 = op(op(e2, e2), e2)) | ~ spl130_287), inference(avatar_component_clause, [], [f2792])).
fof(f2806, plain, (~ (e1 = op(op(op(e2, e2), e2), op(op(e2, e2), e2))) | spl130_289), inference(avatar_component_clause, [], [f2804])).
fof(f10660, plain, (~ spl130_15 | ~ spl130_25), inference(avatar_split_clause, [], [f10657, f1301, f1259])).
fof(f1259, plain, (spl130_15 <=> (e4 = op(e4, e2))), introduced(avatar_definition, [new_symbols(naming, [spl130_15])])).
fof(f10657, plain, (~ (e4 = op(e4, e2)) | ~ spl130_25), inference(backward_demodulation, [], [f685, f1303])).
fof(f1303, plain, ((e4 = op(e4, e0)) | ~ spl130_25), inference(avatar_component_clause, [], [f1301])).
fof(f685, plain, ~ (op(e4, e0) = op(e4, e2)), inference(cnf_transformation, [], [f4])).
fof(f10641, plain, (~ spl130_26 | ~ spl130_36), inference(avatar_split_clause, [], [f10636, f1348, f1306])).
fof(f1306, plain, (spl130_26 <=> (e0 = op(e3, e4))), introduced(avatar_definition, [new_symbols(naming, [spl130_26])])).
fof(f1348, plain, (spl130_36 <=> (e0 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl130_36])])).
fof(f10636, plain, (~ (e0 = op(e3, e4)) | ~ spl130_36), inference(backward_demodulation, [], [f682, f1350])).
fof(f1350, plain, ((e0 = op(e3, e2)) | ~ spl130_36), inference(avatar_component_clause, [], [f1348])).
fof(f10629, plain, (~ spl130_39 | ~ spl130_49), inference(avatar_split_clause, [], [f10625, f1402, f1360])).
fof(f1360, plain, (spl130_39 <=> (e3 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl130_39])])).
fof(f10625, plain, (~ (e3 = op(e3, e2)) | ~ spl130_49), inference(backward_demodulation, [], [f675, f1404])).
fof(f1404, plain, ((e3 = op(e3, e0)) | ~ spl130_49), inference(avatar_component_clause, [], [f1402])).
fof(f675, plain, ~ (op(e3, e0) = op(e3, e2)), inference(cnf_transformation, [], [f4])).
fof(f10619, plain, (~ spl130_55 | ~ spl130_60), inference(avatar_split_clause, [], [f10618, f1448, f1427])).
fof(f1427, plain, (spl130_55 <=> (e4 = op(e2, e4))), introduced(avatar_definition, [new_symbols(naming, [spl130_55])])).
fof(f1448, plain, (spl130_60 <=> (e4 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl130_60])])).
fof(f10618, plain, (~ (e4 = op(e2, e4)) | ~ spl130_60), inference(backward_demodulation, [], [f673, f1450])).
fof(f1450, plain, ((e4 = op(e2, e3)) | ~ spl130_60), inference(avatar_component_clause, [], [f1448])).
fof(f673, plain, ~ (op(e2, e3) = op(e2, e4)), inference(cnf_transformation, [], [f4])).
fof(f10591, plain, (~ spl130_63 | ~ spl130_73), inference(avatar_split_clause, [], [f10586, f1503, f1461])).
fof(f1461, plain, (spl130_63 <=> (e2 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl130_63])])).
fof(f1503, plain, (spl130_73 <=> (e2 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl130_73])])).
fof(f10586, plain, (~ (e2 = op(e2, e2)) | ~ spl130_73), inference(backward_demodulation, [], [f665, f1505])).
fof(f1505, plain, ((e2 = op(e2, e0)) | ~ spl130_73), inference(avatar_component_clause, [], [f1503])).
fof(f10582, plain, (~ spl130_3 | ~ spl130_78), inference(avatar_split_clause, [], [f10579, f1524, f1209])).
fof(f1209, plain, (spl130_3 <=> (e2 = op(e4, e4))), introduced(avatar_definition, [new_symbols(naming, [spl130_3])])).
fof(f10579, plain, (~ (e2 = op(e4, e4)) | ~ spl130_78), inference(backward_demodulation, [], [f640, f1526])).
fof(f1526, plain, ((e2 = op(e1, e4)) | ~ spl130_78), inference(avatar_component_clause, [], [f1524])).
fof(f640, plain, ~ (op(e1, e4) = op(e4, e4)), inference(cnf_transformation, [], [f4])).
fof(f10576, plain, (~ spl130_6 | ~ spl130_81), inference(avatar_split_clause, [], [f10570, f1537, f1222])).
fof(f10570, plain, (~ (e0 = op(e4, e3)) | ~ spl130_81), inference(backward_demodulation, [], [f630, f1539])).
fof(f1539, plain, ((e0 = op(e1, e3)) | ~ spl130_81), inference(avatar_component_clause, [], [f1537])).
fof(f630, plain, ~ (op(e1, e3) = op(e4, e3)), inference(cnf_transformation, [], [f4])).
fof(f10567, plain, (~ spl130_85 | ~ spl130_90), inference(avatar_split_clause, [], [f10564, f1574, f1553])).
fof(f1553, plain, (spl130_85 <=> (e4 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl130_85])])).
fof(f10564, plain, (~ (e4 = op(e1, e3)) | ~ spl130_90), inference(backward_demodulation, [], [f661, f1576])).
fof(f1576, plain, ((e4 = op(e1, e2)) | ~ spl130_90), inference(avatar_component_clause, [], [f1574])).
fof(f661, plain, ~ (op(e1, e2) = op(e1, e3)), inference(cnf_transformation, [], [f4])).
fof(f10566, plain, (~ spl130_15 | ~ spl130_90), inference(avatar_split_clause, [], [f10563, f1574, f1259])).
fof(f10563, plain, (~ (e4 = op(e4, e2)) | ~ spl130_90), inference(backward_demodulation, [], [f620, f1576])).
fof(f620, plain, ~ (op(e1, e2) = op(e4, e2)), inference(cnf_transformation, [], [f4])).
fof(f10560, plain, (spl130_45 | ~ spl130_94 | ~ spl130_286), inference(avatar_split_clause, [], [f10553, f2786, f1591, f1385])).
fof(f1591, plain, (spl130_94 <=> (e3 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl130_94])])).
fof(f2786, plain, (spl130_286 <=> (e4 = op(op(e1, e1), e1))), introduced(avatar_definition, [new_symbols(naming, [spl130_286])])).
fof(f10553, plain, ((e4 = op(e3, e1)) | (~ spl130_94 | ~ spl130_286)), inference(backward_demodulation, [], [f2787, f1593])).
fof(f1593, plain, ((e3 = op(e1, e1)) | ~ spl130_94), inference(avatar_component_clause, [], [f1591])).
fof(f2787, plain, ((e4 = op(op(e1, e1), e1)) | ~ spl130_286), inference(avatar_component_clause, [], [f2786])).
fof(f10558, plain, (~ spl130_43 | ~ spl130_94 | spl130_274), inference(avatar_split_clause, [], [f10550, f2726, f1591, f1377])).
fof(f1377, plain, (spl130_43 <=> (e2 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl130_43])])).
fof(f2726, plain, (spl130_274 <=> (e2 = op(op(e1, e1), e1))), introduced(avatar_definition, [new_symbols(naming, [spl130_274])])).
fof(f10550, plain, (~ (e2 = op(e3, e1)) | (~ spl130_94 | spl130_274)), inference(backward_demodulation, [], [f2728, f1593])).
fof(f2728, plain, (~ (e2 = op(op(e1, e1), e1)) | spl130_274), inference(avatar_component_clause, [], [f2726])).
fof(f10557, plain, (~ spl130_79 | ~ spl130_94), inference(avatar_split_clause, [], [f10548, f1591, f1528])).
fof(f1528, plain, (spl130_79 <=> (e3 = op(e1, e4))), introduced(avatar_definition, [new_symbols(naming, [spl130_79])])).
fof(f10548, plain, (~ (e3 = op(e1, e4)) | ~ spl130_94), inference(backward_demodulation, [], [f660, f1593])).
fof(f660, plain, ~ (op(e1, e1) = op(e1, e4)), inference(cnf_transformation, [], [f4])).
fof(f10542, plain, (~ spl130_87 | ~ spl130_97), inference(avatar_split_clause, [], [f10535, f1604, f1562])).
fof(f1562, plain, (spl130_87 <=> (e1 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl130_87])])).
fof(f1604, plain, (spl130_97 <=> (e1 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl130_97])])).
fof(f10535, plain, (~ (e1 = op(e1, e2)) | ~ spl130_97), inference(backward_demodulation, [], [f655, f1606])).
fof(f1606, plain, ((e1 = op(e1, e0)) | ~ spl130_97), inference(avatar_component_clause, [], [f1604])).
fof(f655, plain, ~ (op(e1, e0) = op(e1, e2)), inference(cnf_transformation, [], [f4])).
fof(f10530, plain, (~ spl130_55 | ~ spl130_105), inference(avatar_split_clause, [], [f10527, f1637, f1427])).
fof(f10527, plain, (~ (e4 = op(e2, e4)) | ~ spl130_105), inference(backward_demodulation, [], [f635, f1639])).
fof(f1639, plain, ((e4 = op(e0, e4)) | ~ spl130_105), inference(avatar_component_clause, [], [f1637])).
fof(f635, plain, ~ (op(e0, e4) = op(e2, e4)), inference(cnf_transformation, [], [f4])).
fof(f10525, plain, (~ spl130_59 | ~ spl130_109), inference(avatar_split_clause, [], [f10520, f1654, f1444])).
fof(f1444, plain, (spl130_59 <=> (e3 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl130_59])])).
fof(f10520, plain, (~ (e3 = op(e2, e3)) | ~ spl130_109), inference(backward_demodulation, [], [f625, f1656])).
fof(f1656, plain, ((e3 = op(e0, e3)) | ~ spl130_109), inference(avatar_component_clause, [], [f1654])).
fof(f625, plain, ~ (op(e0, e3) = op(e2, e3)), inference(cnf_transformation, [], [f4])).
fof(f10517, plain, (~ spl130_63 | ~ spl130_113), inference(avatar_split_clause, [], [f10510, f1671, f1461])).
fof(f1671, plain, (spl130_113 <=> (e2 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl130_113])])).
fof(f10510, plain, (~ (e2 = op(e2, e2)) | ~ spl130_113), inference(backward_demodulation, [], [f615, f1673])).
fof(f1673, plain, ((e2 = op(e0, e2)) | ~ spl130_113), inference(avatar_component_clause, [], [f1671])).
fof(f615, plain, ~ (op(e0, e2) = op(e2, e2)), inference(cnf_transformation, [], [f4])).
fof(f10508, plain, (~ spl130_102 | ~ spl130_117), inference(avatar_split_clause, [], [f10503, f1688, f1625])).
fof(f1625, plain, (spl130_102 <=> (e1 = op(e0, e4))), introduced(avatar_definition, [new_symbols(naming, [spl130_102])])).
fof(f1688, plain, (spl130_117 <=> (e1 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl130_117])])).
fof(f10503, plain, (~ (e1 = op(e0, e4)) | ~ spl130_117), inference(backward_demodulation, [], [f650, f1690])).
fof(f1690, plain, ((e1 = op(e0, e1)) | ~ spl130_117), inference(avatar_component_clause, [], [f1688])).
fof(f650, plain, ~ (op(e0, e1) = op(e0, e4)), inference(cnf_transformation, [], [f4])).
fof(f10507, plain, (~ spl130_67 | ~ spl130_117), inference(avatar_split_clause, [], [f10498, f1688, f1478])).
fof(f1478, plain, (spl130_67 <=> (e1 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl130_67])])).
fof(f10498, plain, (~ (e1 = op(e2, e1)) | ~ spl130_117), inference(backward_demodulation, [], [f605, f1690])).
fof(f10488, plain, (~ spl130_111 | ~ spl130_121), inference(avatar_split_clause, [], [f10475, f1705, f1663])).
fof(f1663, plain, (spl130_111 <=> (e0 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl130_111])])).
fof(f1705, plain, (spl130_121 <=> (e0 = op(e0, e0))), introduced(avatar_definition, [new_symbols(naming, [spl130_121])])).
fof(f10475, plain, (~ (e0 = op(e0, e2)) | ~ spl130_121), inference(backward_demodulation, [], [f645, f1707])).
fof(f1707, plain, ((e0 = op(e0, e0)) | ~ spl130_121), inference(avatar_component_clause, [], [f1705])).
fof(f645, plain, ~ (op(e0, e0) = op(e0, e2)), inference(cnf_transformation, [], [f4])).
fof(f10469, plain, (spl130_25 | ~ spl130_126), inference(avatar_split_clause, [], [f10455, f1726, f1301])).
fof(f1726, plain, (spl130_126 <=> (e0 = unit)), introduced(avatar_definition, [new_symbols(naming, [spl130_126])])).
fof(f10455, plain, ((e4 = op(e4, e0)) | ~ spl130_126), inference(backward_demodulation, [], [f542, f1728])).
fof(f1728, plain, ((e0 = unit) | ~ spl130_126), inference(avatar_component_clause, [], [f1726])).
fof(f10468, plain, (spl130_105 | ~ spl130_126), inference(avatar_split_clause, [], [f10454, f1726, f1637])).
fof(f10454, plain, ((e4 = op(e0, e4)) | ~ spl130_126), inference(backward_demodulation, [], [f541, f1728])).
fof(f10467, plain, (spl130_49 | ~ spl130_126), inference(avatar_split_clause, [], [f10453, f1726, f1402])).
fof(f10453, plain, ((e3 = op(e3, e0)) | ~ spl130_126), inference(backward_demodulation, [], [f540, f1728])).
fof(f10466, plain, (spl130_109 | ~ spl130_126), inference(avatar_split_clause, [], [f10452, f1726, f1654])).
fof(f10452, plain, ((e3 = op(e0, e3)) | ~ spl130_126), inference(backward_demodulation, [], [f539, f1728])).
fof(f10465, plain, (spl130_73 | ~ spl130_126), inference(avatar_split_clause, [], [f10451, f1726, f1503])).
fof(f10451, plain, ((e2 = op(e2, e0)) | ~ spl130_126), inference(backward_demodulation, [], [f538, f1728])).
fof(f10464, plain, (spl130_113 | ~ spl130_126), inference(avatar_split_clause, [], [f10450, f1726, f1671])).
fof(f10450, plain, ((e2 = op(e0, e2)) | ~ spl130_126), inference(backward_demodulation, [], [f537, f1728])).
fof(f10463, plain, (spl130_97 | ~ spl130_126), inference(avatar_split_clause, [], [f10449, f1726, f1604])).
fof(f10449, plain, ((e1 = op(e1, e0)) | ~ spl130_126), inference(backward_demodulation, [], [f536, f1728])).
fof(f10462, plain, (spl130_117 | ~ spl130_126), inference(avatar_split_clause, [], [f10448, f1726, f1688])).
fof(f10448, plain, ((e1 = op(e0, e1)) | ~ spl130_126), inference(backward_demodulation, [], [f535, f1728])).
fof(f535, plain, (e1 = op(unit, e1)), inference(cnf_transformation, [], [f2])).
fof(f10461, plain, (spl130_121 | ~ spl130_126), inference(avatar_split_clause, [], [f10447, f1726, f1705])).
fof(f10447, plain, ((e0 = op(e0, e0)) | ~ spl130_126), inference(backward_demodulation, [], [f534, f1728])).
fof(f10432, plain, (~ spl130_1 | spl130_273 | ~ spl130_286), inference(avatar_split_clause, [], [f10210, f2786, f2722, f1201])).
fof(f2722, plain, (spl130_273 <=> (e0 = op(op(op(e1, e1), e1), op(op(e1, e1), e1)))), introduced(avatar_definition, [new_symbols(naming, [spl130_273])])).
fof(f10210, plain, (~ (e0 = op(e4, e4)) | (spl130_273 | ~ spl130_286)), inference(backward_demodulation, [], [f2724, f2787])).
fof(f2724, plain, (~ (e0 = op(op(op(e1, e1), e1), op(op(e1, e1), e1))) | spl130_273), inference(avatar_component_clause, [], [f2722])).
fof(f10390, plain, (~ spl130_35 | ~ spl130_50), inference(avatar_split_clause, [], [f10388, f1406, f1343])).
fof(f1343, plain, (spl130_35 <=> (e4 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl130_35])])).
fof(f10388, plain, (~ (e4 = op(e3, e3)) | ~ spl130_50), inference(backward_demodulation, [], [f676, f1408])).
fof(f676, plain, ~ (op(e3, e0) = op(e3, e3)), inference(cnf_transformation, [], [f4])).
fof(f10330, plain, (~ spl130_37 | ~ spl130_87), inference(avatar_split_clause, [], [f10323, f1562, f1352])).
fof(f10323, plain, (~ (e1 = op(e3, e2)) | ~ spl130_87), inference(backward_demodulation, [], [f619, f1564])).
fof(f1564, plain, ((e1 = op(e1, e2)) | ~ spl130_87), inference(avatar_component_clause, [], [f1562])).
fof(f619, plain, ~ (op(e1, e2) = op(e3, e2)), inference(cnf_transformation, [], [f4])).
fof(f10318, plain, (~ spl130_81 | ~ spl130_91), inference(avatar_split_clause, [], [f10303, f1579, f1537])).
fof(f1579, plain, (spl130_91 <=> (e0 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl130_91])])).
fof(f10303, plain, (~ (e0 = op(e1, e3)) | ~ spl130_91), inference(backward_demodulation, [], [f659, f1581])).
fof(f1581, plain, ((e0 = op(e1, e1)) | ~ spl130_91), inference(avatar_component_clause, [], [f1579])).
fof(f659, plain, ~ (op(e1, e1) = op(e1, e3)), inference(cnf_transformation, [], [f4])).
fof(f10299, plain, (~ spl130_78 | ~ spl130_98), inference(avatar_split_clause, [], [f10295, f1608, f1524])).
fof(f1608, plain, (spl130_98 <=> (e2 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl130_98])])).
fof(f10295, plain, (~ (e2 = op(e1, e4)) | ~ spl130_98), inference(backward_demodulation, [], [f657, f1610])).
fof(f1610, plain, ((e2 = op(e1, e0)) | ~ spl130_98), inference(avatar_component_clause, [], [f1608])).
fof(f657, plain, ~ (op(e1, e0) = op(e1, e4)), inference(cnf_transformation, [], [f4])).
fof(f10298, plain, (~ spl130_73 | ~ spl130_98), inference(avatar_split_clause, [], [f10289, f1608, f1503])).
fof(f10289, plain, (~ (e2 = op(e2, e0)) | ~ spl130_98), inference(backward_demodulation, [], [f598, f1610])).
fof(f598, plain, ~ (op(e1, e0) = op(e2, e0)), inference(cnf_transformation, [], [f4])).
fof(f10281, plain, (~ spl130_8 | ~ spl130_108), inference(avatar_split_clause, [], [f10277, f1650, f1230])).
fof(f1230, plain, (spl130_8 <=> (e2 = op(e4, e3))), introduced(avatar_definition, [new_symbols(naming, [spl130_8])])).
fof(f1650, plain, (spl130_108 <=> (e2 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl130_108])])).
fof(f10277, plain, (~ (e2 = op(e4, e3)) | ~ spl130_108), inference(backward_demodulation, [], [f627, f1652])).
fof(f1652, plain, ((e2 = op(e0, e3)) | ~ spl130_108), inference(avatar_component_clause, [], [f1650])).
fof(f10262, plain, (~ spl130_105 | ~ spl130_120), inference(avatar_split_clause, [], [f10260, f1700, f1637])).
fof(f1700, plain, (spl130_120 <=> (e4 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl130_120])])).
fof(f10260, plain, (~ (e4 = op(e0, e4)) | ~ spl130_120), inference(backward_demodulation, [], [f650, f1702])).
fof(f1702, plain, ((e4 = op(e0, e1)) | ~ spl130_120), inference(avatar_component_clause, [], [f1700])).
fof(f10261, plain, (~ spl130_70 | ~ spl130_120), inference(avatar_split_clause, [], [f10256, f1700, f1490])).
fof(f1490, plain, (spl130_70 <=> (e4 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl130_70])])).
fof(f10256, plain, (~ (e4 = op(e2, e1)) | ~ spl130_120), inference(backward_demodulation, [], [f605, f1702])).
fof(f10235, plain, (spl130_15 | ~ spl130_128), inference(avatar_split_clause, [], [f10223, f1734, f1259])).
fof(f1734, plain, (spl130_128 <=> (e2 = unit)), introduced(avatar_definition, [new_symbols(naming, [spl130_128])])).
fof(f10223, plain, ((e4 = op(e4, e2)) | ~ spl130_128), inference(backward_demodulation, [], [f542, f1736])).
fof(f1736, plain, ((e2 = unit) | ~ spl130_128), inference(avatar_component_clause, [], [f1734])).
fof(f10234, plain, (spl130_55 | ~ spl130_128), inference(avatar_split_clause, [], [f10222, f1734, f1427])).
fof(f10222, plain, ((e4 = op(e2, e4)) | ~ spl130_128), inference(backward_demodulation, [], [f541, f1736])).
fof(f10233, plain, (spl130_39 | ~ spl130_128), inference(avatar_split_clause, [], [f10221, f1734, f1360])).
fof(f10221, plain, ((e3 = op(e3, e2)) | ~ spl130_128), inference(backward_demodulation, [], [f540, f1736])).
fof(f10232, plain, (spl130_59 | ~ spl130_128), inference(avatar_split_clause, [], [f10220, f1734, f1444])).
fof(f10220, plain, ((e3 = op(e2, e3)) | ~ spl130_128), inference(backward_demodulation, [], [f539, f1736])).
fof(f10231, plain, (spl130_63 | ~ spl130_128), inference(avatar_split_clause, [], [f10219, f1734, f1461])).
fof(f10219, plain, ((e2 = op(e2, e2)) | ~ spl130_128), inference(backward_demodulation, [], [f538, f1736])).
fof(f10229, plain, (spl130_87 | ~ spl130_128), inference(avatar_split_clause, [], [f10217, f1734, f1562])).
fof(f10217, plain, ((e1 = op(e1, e2)) | ~ spl130_128), inference(backward_demodulation, [], [f536, f1736])).
fof(f10228, plain, (spl130_67 | ~ spl130_128), inference(avatar_split_clause, [], [f10216, f1734, f1478])).
fof(f10216, plain, ((e1 = op(e2, e1)) | ~ spl130_128), inference(backward_demodulation, [], [f535, f1736])).
fof(f10227, plain, (spl130_111 | ~ spl130_128), inference(avatar_split_clause, [], [f10215, f1734, f1663])).
fof(f10215, plain, ((e0 = op(e0, e2)) | ~ spl130_128), inference(backward_demodulation, [], [f534, f1736])).
fof(f10226, plain, (spl130_71 | ~ spl130_128), inference(avatar_split_clause, [], [f10214, f1734, f1495])).
fof(f1495, plain, (spl130_71 <=> (e0 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl130_71])])).
fof(f10214, plain, ((e0 = op(e2, e0)) | ~ spl130_128), inference(backward_demodulation, [], [f533, f1736])).
fof(f10208, plain, (~ spl130_3 | ~ spl130_288 | spl130_307), inference(avatar_split_clause, [], [f10205, f2902, f2798, f1209])).
fof(f2798, plain, (spl130_288 <=> (e4 = op(op(e3, e3), e3))), introduced(avatar_definition, [new_symbols(naming, [spl130_288])])).
fof(f2902, plain, (spl130_307 <=> (e2 = op(op(op(e3, e3), e3), op(op(e3, e3), e3)))), introduced(avatar_definition, [new_symbols(naming, [spl130_307])])).
fof(f10205, plain, (~ (e2 = op(e4, e4)) | (~ spl130_288 | spl130_307)), inference(backward_demodulation, [], [f2904, f2799])).
fof(f2799, plain, ((e4 = op(op(e3, e3), e3)) | ~ spl130_288), inference(avatar_component_clause, [], [f2798])).
fof(f2904, plain, (~ (e2 = op(op(op(e3, e3), e3), op(op(e3, e3), e3))) | spl130_307), inference(avatar_component_clause, [], [f2902])).
fof(f10199, plain, (~ spl130_3 | ~ spl130_304 | spl130_309), inference(avatar_split_clause, [], [f10198, f2914, f2882, f1209])).
fof(f2914, plain, (spl130_309 <=> (e2 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), introduced(avatar_definition, [new_symbols(naming, [spl130_309])])).
fof(f10198, plain, (~ (e2 = op(e4, e4)) | (~ spl130_304 | spl130_309)), inference(backward_demodulation, [], [f2916, f2883])).
fof(f2916, plain, (~ (e2 = op(op(op(e0, e0), e0), op(op(e0, e0), e0))) | spl130_309), inference(avatar_component_clause, [], [f2914])).
fof(f10180, plain, (~ spl130_2 | ~ spl130_94 | spl130_271), inference(avatar_contradiction_clause, [], [f10179])).
fof(f10179, plain, ($false | (~ spl130_2 | ~ spl130_94 | spl130_271)), inference(subsumption_resolution, [], [f10174, f1593])).
fof(f10174, plain, (~ (e3 = op(e1, e1)) | (~ spl130_2 | spl130_271)), inference(backward_demodulation, [], [f2714, f1207])).
fof(f1207, plain, ((e1 = op(e4, e4)) | ~ spl130_2), inference(avatar_component_clause, [], [f1205])).
fof(f2714, plain, (~ (e3 = op(op(e4, e4), op(e4, e4))) | spl130_271), inference(avatar_component_clause, [], [f2712])).
fof(f2712, plain, (spl130_271 <=> (e3 = op(op(e4, e4), op(e4, e4)))), introduced(avatar_definition, [new_symbols(naming, [spl130_271])])).
fof(f10169, plain, (~ spl130_1 | ~ spl130_16), inference(avatar_split_clause, [], [f10164, f1264, f1201])).
fof(f1264, plain, (spl130_16 <=> (e0 = op(e4, e1))), introduced(avatar_definition, [new_symbols(naming, [spl130_16])])).
fof(f10164, plain, (~ (e0 = op(e4, e4)) | ~ spl130_16), inference(backward_demodulation, [], [f690, f1266])).
fof(f1266, plain, ((e0 = op(e4, e1)) | ~ spl130_16), inference(avatar_component_clause, [], [f1264])).
fof(f690, plain, ~ (op(e4, e1) = op(e4, e4)), inference(cnf_transformation, [], [f4])).
fof(f10163, plain, (~ spl130_20 | ~ spl130_25), inference(avatar_split_clause, [], [f10161, f1301, f1280])).
fof(f10161, plain, (~ (e4 = op(e4, e1)) | ~ spl130_25), inference(backward_demodulation, [], [f684, f1303])).
fof(f684, plain, ~ (op(e4, e0) = op(e4, e1)), inference(cnf_transformation, [], [f4])).
fof(f10147, plain, (~ spl130_32 | ~ spl130_37), inference(avatar_split_clause, [], [f10142, f1352, f1331])).
fof(f1331, plain, (spl130_32 <=> (e1 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl130_32])])).
fof(f10142, plain, (~ (e1 = op(e3, e3)) | ~ spl130_37), inference(backward_demodulation, [], [f681, f1354])).
fof(f681, plain, ~ (op(e3, e2) = op(e3, e3)), inference(cnf_transformation, [], [f4])).
fof(f10141, plain, (~ spl130_28 | ~ spl130_43), inference(avatar_split_clause, [], [f10138, f1377, f1314])).
fof(f1314, plain, (spl130_28 <=> (e2 = op(e3, e4))), introduced(avatar_definition, [new_symbols(naming, [spl130_28])])).
fof(f10138, plain, (~ (e2 = op(e3, e4)) | ~ spl130_43), inference(backward_demodulation, [], [f680, f1379])).
fof(f1379, plain, ((e2 = op(e3, e1)) | ~ spl130_43), inference(avatar_component_clause, [], [f1377])).
fof(f680, plain, ~ (op(e3, e1) = op(e3, e4)), inference(cnf_transformation, [], [f4])).
fof(f10134, plain, (~ spl130_44 | ~ spl130_49), inference(avatar_split_clause, [], [f10129, f1402, f1381])).
fof(f10129, plain, (~ (e3 = op(e3, e1)) | ~ spl130_49), inference(backward_demodulation, [], [f674, f1404])).
fof(f10124, plain, (~ spl130_52 | ~ spl130_57), inference(avatar_split_clause, [], [f10119, f1436, f1415])).
fof(f1415, plain, (spl130_52 <=> (e1 = op(e2, e4))), introduced(avatar_definition, [new_symbols(naming, [spl130_52])])).
fof(f1436, plain, (spl130_57 <=> (e1 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl130_57])])).
fof(f10119, plain, (~ (e1 = op(e2, e4)) | ~ spl130_57), inference(backward_demodulation, [], [f673, f1438])).
fof(f1438, plain, ((e1 = op(e2, e3)) | ~ spl130_57), inference(avatar_component_clause, [], [f1436])).
fof(f10112, plain, (~ spl130_56 | ~ spl130_61), inference(avatar_split_clause, [], [f10100, f1453, f1432])).
fof(f1432, plain, (spl130_56 <=> (e0 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl130_56])])).
fof(f1453, plain, (spl130_61 <=> (e0 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl130_61])])).
fof(f10100, plain, (~ (e0 = op(e2, e3)) | ~ spl130_61), inference(backward_demodulation, [], [f671, f1455])).
fof(f1455, plain, ((e0 = op(e2, e2)) | ~ spl130_61), inference(avatar_component_clause, [], [f1453])).
fof(f671, plain, ~ (op(e2, e2) = op(e2, e3)), inference(cnf_transformation, [], [f4])).
fof(f10098, plain, (~ spl130_65 | ~ spl130_70), inference(avatar_split_clause, [], [f10094, f1490, f1469])).
fof(f1469, plain, (spl130_65 <=> (e4 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl130_65])])).
fof(f10094, plain, (~ (e4 = op(e2, e2)) | ~ spl130_70), inference(backward_demodulation, [], [f668, f1492])).
fof(f1492, plain, ((e4 = op(e2, e1)) | ~ spl130_70), inference(avatar_component_clause, [], [f1490])).
fof(f668, plain, ~ (op(e2, e1) = op(e2, e2)), inference(cnf_transformation, [], [f4])).
fof(f10097, plain, (~ spl130_20 | ~ spl130_70), inference(avatar_split_clause, [], [f10093, f1490, f1280])).
fof(f10093, plain, (~ (e4 = op(e4, e1)) | ~ spl130_70), inference(backward_demodulation, [], [f612, f1492])).
fof(f612, plain, ~ (op(e2, e1) = op(e4, e1)), inference(cnf_transformation, [], [f4])).
fof(f10091, plain, (~ spl130_68 | ~ spl130_73), inference(avatar_split_clause, [], [f10085, f1503, f1482])).
fof(f10085, plain, (~ (e2 = op(e2, e1)) | ~ spl130_73), inference(backward_demodulation, [], [f664, f1505])).
fof(f664, plain, ~ (op(e2, e0) = op(e2, e1)), inference(cnf_transformation, [], [f4])).
fof(f10082, plain, (~ spl130_28 | ~ spl130_78), inference(avatar_split_clause, [], [f10078, f1524, f1314])).
fof(f10078, plain, (~ (e2 = op(e3, e4)) | ~ spl130_78), inference(backward_demodulation, [], [f639, f1526])).
fof(f639, plain, ~ (op(e1, e4) = op(e3, e4)), inference(cnf_transformation, [], [f4])).
fof(f10076, plain, (~ spl130_56 | ~ spl130_81), inference(avatar_split_clause, [], [f10069, f1537, f1432])).
fof(f10069, plain, (~ (e0 = op(e2, e3)) | ~ spl130_81), inference(backward_demodulation, [], [f628, f1539])).
fof(f628, plain, ~ (op(e1, e3) = op(e2, e3)), inference(cnf_transformation, [], [f4])).
fof(f10068, plain, (~ spl130_80 | ~ spl130_90), inference(avatar_split_clause, [], [f10066, f1574, f1532])).
fof(f10066, plain, (~ (e4 = op(e1, e4)) | ~ spl130_90), inference(backward_demodulation, [], [f662, f1576])).
fof(f662, plain, ~ (op(e1, e2) = op(e1, e4)), inference(cnf_transformation, [], [f4])).
fof(f10067, plain, (~ spl130_65 | ~ spl130_90), inference(avatar_split_clause, [], [f10063, f1574, f1469])).
fof(f10063, plain, (~ (e4 = op(e2, e2)) | ~ spl130_90), inference(backward_demodulation, [], [f618, f1576])).
fof(f618, plain, ~ (op(e1, e2) = op(e2, e2)), inference(cnf_transformation, [], [f4])).
fof(f10046, plain, (~ spl130_92 | ~ spl130_97), inference(avatar_split_clause, [], [f10038, f1604, f1583])).
fof(f10038, plain, (~ (e1 = op(e1, e1)) | ~ spl130_97), inference(backward_demodulation, [], [f654, f1606])).
fof(f654, plain, ~ (op(e1, e0) = op(e1, e1)), inference(cnf_transformation, [], [f4])).
fof(f10034, plain, (~ spl130_80 | ~ spl130_105), inference(avatar_split_clause, [], [f10030, f1637, f1532])).
fof(f10030, plain, (~ (e4 = op(e1, e4)) | ~ spl130_105), inference(backward_demodulation, [], [f634, f1639])).
fof(f10028, plain, (~ spl130_84 | ~ spl130_109), inference(avatar_split_clause, [], [f10023, f1654, f1549])).
fof(f10023, plain, (~ (e3 = op(e1, e3)) | ~ spl130_109), inference(backward_demodulation, [], [f624, f1656])).
fof(f624, plain, ~ (op(e0, e3) = op(e1, e3)), inference(cnf_transformation, [], [f4])).
fof(f10022, plain, (~ spl130_88 | ~ spl130_113), inference(avatar_split_clause, [], [f10015, f1671, f1566])).
fof(f10015, plain, (~ (e2 = op(e1, e2)) | ~ spl130_113), inference(backward_demodulation, [], [f614, f1673])).
fof(f10014, plain, (~ spl130_112 | ~ spl130_117), inference(avatar_split_clause, [], [f10007, f1688, f1667])).
fof(f1667, plain, (spl130_112 <=> (e1 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl130_112])])).
fof(f10007, plain, (~ (e1 = op(e0, e2)) | ~ spl130_117), inference(backward_demodulation, [], [f648, f1690])).
fof(f648, plain, ~ (op(e0, e1) = op(e0, e2)), inference(cnf_transformation, [], [f4])).
fof(f9993, plain, (~ spl130_116 | ~ spl130_121), inference(avatar_split_clause, [], [f9979, f1705, f1684])).
fof(f9979, plain, (~ (e0 = op(e0, e1)) | ~ spl130_121), inference(backward_demodulation, [], [f644, f1707])).
fof(f644, plain, ~ (op(e0, e0) = op(e0, e1)), inference(cnf_transformation, [], [f4])).
fof(f9948, plain, (~ spl130_61 | spl130_273 | ~ spl130_274), inference(avatar_split_clause, [], [f9947, f2726, f2722, f1453])).
fof(f9947, plain, (~ (e0 = op(e2, e2)) | (spl130_273 | ~ spl130_274)), inference(forward_demodulation, [], [f2724, f2727])).
fof(f2727, plain, ((e2 = op(op(e1, e1), e1)) | ~ spl130_274), inference(avatar_component_clause, [], [f2726])).
fof(f9943, plain, (~ spl130_32 | ~ spl130_283 | spl130_289), inference(avatar_split_clause, [], [f9942, f2804, f2770, f1331])).
fof(f2770, plain, (spl130_283 <=> (e3 = op(op(e2, e2), e2))), introduced(avatar_definition, [new_symbols(naming, [spl130_283])])).
fof(f9942, plain, (~ (e1 = op(e3, e3)) | (~ spl130_283 | spl130_289)), inference(forward_demodulation, [], [f2806, f2771])).
fof(f2771, plain, ((e3 = op(op(e2, e2), e2)) | ~ spl130_283), inference(avatar_component_clause, [], [f2770])).
fof(f9936, plain, (~ spl130_58 | ~ spl130_8), inference(avatar_split_clause, [], [f9935, f1230, f1440])).
fof(f1440, plain, (spl130_58 <=> (e2 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl130_58])])).
fof(f9935, plain, (~ (e2 = op(e2, e3)) | ~ spl130_8), inference(forward_demodulation, [], [f632, f1232])).
fof(f1232, plain, ((e2 = op(e4, e3)) | ~ spl130_8), inference(avatar_component_clause, [], [f1230])).
fof(f632, plain, ~ (op(e2, e3) = op(e4, e3)), inference(cnf_transformation, [], [f4])).
fof(f9918, plain, (~ spl130_27 | ~ spl130_32), inference(avatar_split_clause, [], [f9911, f1331, f1310])).
fof(f9911, plain, (~ (e1 = op(e3, e4)) | ~ spl130_32), inference(backward_demodulation, [], [f683, f1333])).
fof(f1333, plain, ((e1 = op(e3, e3)) | ~ spl130_32), inference(avatar_component_clause, [], [f1331])).
fof(f683, plain, ~ (op(e3, e3) = op(e3, e4)), inference(cnf_transformation, [], [f4])).
fof(f9893, plain, (~ spl130_54 | ~ spl130_74), inference(avatar_split_clause, [], [f9891, f1507, f1423])).
fof(f1423, plain, (spl130_54 <=> (e3 = op(e2, e4))), introduced(avatar_definition, [new_symbols(naming, [spl130_54])])).
fof(f1507, plain, (spl130_74 <=> (e3 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl130_74])])).
fof(f9891, plain, (~ (e3 = op(e2, e4)) | ~ spl130_74), inference(backward_demodulation, [], [f667, f1509])).
fof(f1509, plain, ((e3 = op(e2, e0)) | ~ spl130_74), inference(avatar_component_clause, [], [f1507])).
fof(f667, plain, ~ (op(e2, e0) = op(e2, e4)), inference(cnf_transformation, [], [f4])).
fof(f9888, plain, (~ spl130_32 | ~ spl130_1 | ~ spl130_104 | spl130_293), inference(avatar_split_clause, [], [f9883, f2824, f1633, f1201, f1331])).
fof(f1633, plain, (spl130_104 <=> (e3 = op(e0, e4))), introduced(avatar_definition, [new_symbols(naming, [spl130_104])])).
fof(f2824, plain, (spl130_293 <=> (e1 = op(op(op(e4, e4), e4), op(op(e4, e4), e4)))), introduced(avatar_definition, [new_symbols(naming, [spl130_293])])).
fof(f9883, plain, (~ (e1 = op(e3, e3)) | (~ spl130_1 | ~ spl130_104 | spl130_293)), inference(backward_demodulation, [], [f9516, f1635])).
fof(f1635, plain, ((e3 = op(e0, e4)) | ~ spl130_104), inference(avatar_component_clause, [], [f1633])).
fof(f9516, plain, (~ (e1 = op(op(e0, e4), op(e0, e4))) | (~ spl130_1 | spl130_293)), inference(backward_demodulation, [], [f2826, f1203])).
fof(f1203, plain, ((e0 = op(e4, e4)) | ~ spl130_1), inference(avatar_component_clause, [], [f1201])).
fof(f2826, plain, (~ (e1 = op(op(op(e4, e4), e4), op(op(e4, e4), e4))) | spl130_293), inference(avatar_component_clause, [], [f2824])).
fof(f9887, plain, (~ spl130_54 | ~ spl130_104), inference(avatar_split_clause, [], [f9880, f1633, f1423])).
fof(f9880, plain, (~ (e3 = op(e2, e4)) | ~ spl130_104), inference(backward_demodulation, [], [f635, f1635])).
fof(f9876, plain, (~ spl130_107 | ~ spl130_112), inference(avatar_split_clause, [], [f9870, f1667, f1646])).
fof(f9870, plain, (~ (e1 = op(e0, e3)) | ~ spl130_112), inference(backward_demodulation, [], [f651, f1669])).
fof(f1669, plain, ((e1 = op(e0, e2)) | ~ spl130_112), inference(avatar_component_clause, [], [f1667])).
fof(f651, plain, ~ (op(e0, e2) = op(e0, e3)), inference(cnf_transformation, [], [f4])).
fof(f9868, plain, (~ spl130_62 | ~ spl130_123 | spl130_311), inference(avatar_split_clause, [], [f9862, f2929, f1713, f1457])).
fof(f1713, plain, (spl130_123 <=> (op(e0, e0) = e2)), introduced(avatar_definition, [new_symbols(naming, [spl130_123])])).
fof(f2929, plain, (spl130_311 <=> (e1 = op(op(e0, e0), op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl130_311])])).
fof(f9862, plain, (~ (e1 = op(e2, e2)) | (~ spl130_123 | spl130_311)), inference(backward_demodulation, [], [f2931, f1715])).
fof(f1715, plain, ((op(e0, e0) = e2) | ~ spl130_123), inference(avatar_component_clause, [], [f1713])).
fof(f2931, plain, (~ (e1 = op(op(e0, e0), op(e0, e0))) | spl130_311), inference(avatar_component_clause, [], [f2929])).
fof(f9866, plain, (spl130_73 | ~ spl130_123 | ~ spl130_296), inference(avatar_contradiction_clause, [], [f9865])).
fof(f9865, plain, ($false | (spl130_73 | ~ spl130_123 | ~ spl130_296)), inference(subsumption_resolution, [], [f9857, f1504])).
fof(f1504, plain, (~ (e2 = op(e2, e0)) | spl130_73), inference(avatar_component_clause, [], [f1503])).
fof(f9857, plain, ((e2 = op(e2, e0)) | (~ spl130_123 | ~ spl130_296)), inference(backward_demodulation, [], [f2839, f1715])).
fof(f2839, plain, ((e2 = op(op(e0, e0), e0)) | ~ spl130_296), inference(avatar_component_clause, [], [f2838])).
fof(f2838, plain, (spl130_296 <=> (e2 = op(op(e0, e0), e0))), introduced(avatar_definition, [new_symbols(naming, [spl130_296])])).
fof(f9864, plain, (~ spl130_103 | ~ spl130_123), inference(avatar_split_clause, [], [f9854, f1713, f1629])).
fof(f9854, plain, (~ (e2 = op(e0, e4)) | ~ spl130_123), inference(backward_demodulation, [], [f647, f1715])).
fof(f647, plain, ~ (op(e0, e0) = op(e0, e4)), inference(cnf_transformation, [], [f4])).
fof(f9863, plain, (~ spl130_48 | ~ spl130_123), inference(avatar_split_clause, [], [f9851, f1713, f1398])).
fof(f1398, plain, (spl130_48 <=> (e2 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl130_48])])).
fof(f9851, plain, (~ (e2 = op(e3, e0)) | ~ spl130_123), inference(backward_demodulation, [], [f596, f1715])).
fof(f9820, plain, (~ spl130_65 | ~ spl130_296 | spl130_320), inference(avatar_split_clause, [], [f9596, f3006, f2838, f1469])).
fof(f3006, plain, (spl130_320 <=> (e4 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), introduced(avatar_definition, [new_symbols(naming, [spl130_320])])).
fof(f9596, plain, (~ (e4 = op(e2, e2)) | (~ spl130_296 | spl130_320)), inference(backward_demodulation, [], [f3008, f2839])).
fof(f3008, plain, (~ (e4 = op(op(op(e0, e0), e0), op(op(e0, e0), e0))) | spl130_320), inference(avatar_component_clause, [], [f3006])).
fof(f9819, plain, (~ spl130_65 | ~ spl130_277 | spl130_319), inference(avatar_split_clause, [], [f9607, f3000, f2740, f1469])).
fof(f2740, plain, (spl130_277 <=> (e2 = op(op(e3, e3), e3))), introduced(avatar_definition, [new_symbols(naming, [spl130_277])])).
fof(f3000, plain, (spl130_319 <=> (e4 = op(op(op(e3, e3), e3), op(op(e3, e3), e3)))), introduced(avatar_definition, [new_symbols(naming, [spl130_319])])).
fof(f9607, plain, (~ (e4 = op(e2, e2)) | (~ spl130_277 | spl130_319)), inference(backward_demodulation, [], [f3002, f2741])).
fof(f2741, plain, ((e2 = op(op(e3, e3), e3)) | ~ spl130_277), inference(avatar_component_clause, [], [f2740])).
fof(f3002, plain, (~ (e4 = op(op(op(e3, e3), e3), op(op(e3, e3), e3))) | spl130_319), inference(avatar_component_clause, [], [f3000])).
fof(f9811, plain, (~ spl130_12 | ~ spl130_22), inference(avatar_split_clause, [], [f9807, f1289, f1247])).
fof(f1247, plain, (spl130_12 <=> (e1 = op(e4, e2))), introduced(avatar_definition, [new_symbols(naming, [spl130_12])])).
fof(f1289, plain, (spl130_22 <=> (e1 = op(e4, e0))), introduced(avatar_definition, [new_symbols(naming, [spl130_22])])).
fof(f9807, plain, (~ (e1 = op(e4, e2)) | ~ spl130_22), inference(backward_demodulation, [], [f685, f1291])).
fof(f1291, plain, ((e1 = op(e4, e0)) | ~ spl130_22), inference(avatar_component_clause, [], [f1289])).
fof(f9785, plain, (~ spl130_28 | ~ spl130_48), inference(avatar_split_clause, [], [f9782, f1398, f1314])).
fof(f9782, plain, (~ (e2 = op(e3, e4)) | ~ spl130_48), inference(backward_demodulation, [], [f677, f1400])).
fof(f1400, plain, ((e2 = op(e3, e0)) | ~ spl130_48), inference(avatar_component_clause, [], [f1398])).
fof(f677, plain, ~ (op(e3, e0) = op(e3, e4)), inference(cnf_transformation, [], [f4])).
fof(f9767, plain, (~ spl130_52 | ~ spl130_62), inference(avatar_split_clause, [], [f9755, f1457, f1415])).
fof(f9755, plain, (~ (e1 = op(e2, e4)) | ~ spl130_62), inference(backward_demodulation, [], [f672, f1459])).
fof(f1459, plain, ((e1 = op(e2, e2)) | ~ spl130_62), inference(avatar_component_clause, [], [f1457])).
fof(f672, plain, ~ (op(e2, e2) = op(e2, e4)), inference(cnf_transformation, [], [f4])).
fof(f9746, plain, (~ spl130_70 | ~ spl130_75), inference(avatar_split_clause, [], [f9742, f1511, f1490])).
fof(f1511, plain, (spl130_75 <=> (e4 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl130_75])])).
fof(f9742, plain, (~ (e4 = op(e2, e1)) | ~ spl130_75), inference(backward_demodulation, [], [f664, f1513])).
fof(f1513, plain, ((e4 = op(e2, e0)) | ~ spl130_75), inference(avatar_component_clause, [], [f1511])).
fof(f9745, plain, (~ spl130_25 | ~ spl130_75), inference(avatar_split_clause, [], [f9741, f1511, f1301])).
fof(f9741, plain, (~ (e4 = op(e4, e0)) | ~ spl130_75), inference(backward_demodulation, [], [f602, f1513])).
fof(f602, plain, ~ (op(e2, e0) = op(e4, e0)), inference(cnf_transformation, [], [f4])).
fof(f9702, plain, (~ spl130_86 | ~ spl130_96), inference(avatar_split_clause, [], [f9695, f1600, f1558])).
fof(f1558, plain, (spl130_86 <=> (e0 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl130_86])])).
fof(f9695, plain, (~ (e0 = op(e1, e2)) | ~ spl130_96), inference(backward_demodulation, [], [f655, f1602])).
fof(f9688, plain, (~ spl130_28 | ~ spl130_103), inference(avatar_split_clause, [], [f9684, f1629, f1314])).
fof(f9684, plain, (~ (e2 = op(e3, e4)) | ~ spl130_103), inference(backward_demodulation, [], [f636, f1631])).
fof(f636, plain, ~ (op(e0, e4) = op(e3, e4)), inference(cnf_transformation, [], [f4])).
fof(f9681, plain, (~ spl130_32 | ~ spl130_107), inference(avatar_split_clause, [], [f9676, f1646, f1331])).
fof(f9676, plain, (~ (e1 = op(e3, e3)) | ~ spl130_107), inference(backward_demodulation, [], [f626, f1648])).
fof(f626, plain, ~ (op(e0, e3) = op(e3, e3)), inference(cnf_transformation, [], [f4])).
fof(f9654, plain, (~ spl130_32 | ~ spl130_124 | spl130_311), inference(avatar_split_clause, [], [f9649, f2929, f1717, f1331])).
fof(f9649, plain, (~ (e1 = op(e3, e3)) | (~ spl130_124 | spl130_311)), inference(backward_demodulation, [], [f2931, f1719])).
fof(f9652, plain, (spl130_48 | ~ spl130_124 | ~ spl130_296), inference(avatar_split_clause, [], [f9644, f2838, f1717, f1398])).
fof(f9644, plain, ((e2 = op(e3, e0)) | (~ spl130_124 | ~ spl130_296)), inference(backward_demodulation, [], [f2839, f1719])).
fof(f9600, plain, (~ spl130_64 | ~ spl130_296 | spl130_316), inference(avatar_split_clause, [], [f9597, f2966, f2838, f1465])).
fof(f1465, plain, (spl130_64 <=> (e3 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl130_64])])).
fof(f2966, plain, (spl130_316 <=> (e3 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), introduced(avatar_definition, [new_symbols(naming, [spl130_316])])).
fof(f9597, plain, (~ (e3 = op(e2, e2)) | (~ spl130_296 | spl130_316)), inference(backward_demodulation, [], [f2968, f2839])).
fof(f2968, plain, (~ (e3 = op(op(op(e0, e0), e0), op(op(e0, e0), e0))) | spl130_316), inference(avatar_component_clause, [], [f2966])).
fof(f9599, plain, (~ spl130_62 | spl130_295 | ~ spl130_296), inference(avatar_split_clause, [], [f9594, f2838, f2834, f1457])).
fof(f2834, plain, (spl130_295 <=> (e1 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), introduced(avatar_definition, [new_symbols(naming, [spl130_295])])).
fof(f9594, plain, (~ (e1 = op(e2, e2)) | (spl130_295 | ~ spl130_296)), inference(backward_demodulation, [], [f2836, f2839])).
fof(f2836, plain, (~ (e1 = op(op(op(e0, e0), e0), op(op(e0, e0), e0))) | spl130_295), inference(avatar_component_clause, [], [f2834])).
fof(f9591, plain, (~ spl130_123 | ~ spl130_1 | spl130_272), inference(avatar_split_clause, [], [f9513, f2717, f1201, f1713])).
fof(f2717, plain, (spl130_272 <=> (e2 = op(op(e4, e4), op(e4, e4)))), introduced(avatar_definition, [new_symbols(naming, [spl130_272])])).
fof(f9513, plain, (~ (op(e0, e0) = e2) | (~ spl130_1 | spl130_272)), inference(backward_demodulation, [], [f2719, f1203])).
fof(f2719, plain, (~ (e2 = op(op(e4, e4), op(e4, e4))) | spl130_272), inference(avatar_component_clause, [], [f2717])).
fof(f9561, plain, (~ spl130_1 | ~ spl130_70 | ~ spl130_93 | spl130_273), inference(avatar_contradiction_clause, [], [f9560])).
fof(f9560, plain, ($false | (~ spl130_1 | ~ spl130_70 | ~ spl130_93 | spl130_273)), inference(subsumption_resolution, [], [f9559, f1203])).
fof(f9559, plain, (~ (e0 = op(e4, e4)) | (~ spl130_70 | ~ spl130_93 | spl130_273)), inference(forward_demodulation, [], [f9558, f1492])).
fof(f9558, plain, (~ (e0 = op(op(e2, e1), op(e2, e1))) | (~ spl130_93 | spl130_273)), inference(forward_demodulation, [], [f2724, f1589])).
fof(f1589, plain, ((e2 = op(e1, e1)) | ~ spl130_93), inference(avatar_component_clause, [], [f1587])).
fof(f1587, plain, (spl130_93 <=> (e2 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl130_93])])).
fof(f9556, plain, (~ spl130_32 | ~ spl130_64 | spl130_284), inference(avatar_contradiction_clause, [], [f9555])).
fof(f9555, plain, ($false | (~ spl130_32 | ~ spl130_64 | spl130_284)), inference(subsumption_resolution, [], [f9554, f1333])).
fof(f9554, plain, (~ (e1 = op(e3, e3)) | (~ spl130_64 | spl130_284)), inference(forward_demodulation, [], [f2777, f1467])).
fof(f1467, plain, ((e3 = op(e2, e2)) | ~ spl130_64), inference(avatar_component_clause, [], [f1465])).
fof(f2777, plain, (~ (e1 = op(op(e2, e2), op(e2, e2))) | spl130_284), inference(avatar_component_clause, [], [f2775])).
fof(f2775, plain, (spl130_284 <=> (e1 = op(op(e2, e2), op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl130_284])])).
fof(f9501, plain, (~ spl130_7 | ~ spl130_12), inference(avatar_split_clause, [], [f9496, f1247, f1226])).
fof(f9496, plain, (~ (e1 = op(e4, e3)) | ~ spl130_12), inference(backward_demodulation, [], [f691, f1249])).
fof(f1249, plain, ((e1 = op(e4, e2)) | ~ spl130_12), inference(avatar_component_clause, [], [f1247])).
fof(f691, plain, ~ (op(e4, e2) = op(e4, e3)), inference(cnf_transformation, [], [f4])).
fof(f9495, plain, (~ spl130_3 | ~ spl130_28), inference(avatar_split_clause, [], [f9492, f1314, f1209])).
fof(f9492, plain, (~ (e2 = op(e4, e4)) | ~ spl130_28), inference(backward_demodulation, [], [f643, f1316])).
fof(f1316, plain, ((e2 = op(e3, e4)) | ~ spl130_28), inference(avatar_component_clause, [], [f1314])).
fof(f643, plain, ~ (op(e3, e4) = op(e4, e4)), inference(cnf_transformation, [], [f4])).
fof(f9461, plain, (~ spl130_11 | ~ spl130_86), inference(avatar_split_clause, [], [f9455, f1558, f1243])).
fof(f1243, plain, (spl130_11 <=> (e0 = op(e4, e2))), introduced(avatar_definition, [new_symbols(naming, [spl130_11])])).
fof(f9455, plain, (~ (e0 = op(e4, e2)) | ~ spl130_86), inference(backward_demodulation, [], [f620, f1560])).
fof(f1560, plain, ((e0 = op(e1, e2)) | ~ spl130_86), inference(avatar_component_clause, [], [f1558])).
fof(f9453, plain, (~ spl130_64 | ~ spl130_93 | spl130_276), inference(avatar_contradiction_clause, [], [f9452])).
fof(f9452, plain, ($false | (~ spl130_64 | ~ spl130_93 | spl130_276)), inference(subsumption_resolution, [], [f9443, f1467])).
fof(f9443, plain, (~ (e3 = op(e2, e2)) | (~ spl130_93 | spl130_276)), inference(backward_demodulation, [], [f2737, f1589])).
fof(f2737, plain, (~ (e3 = op(op(e1, e1), op(e1, e1))) | spl130_276), inference(avatar_component_clause, [], [f2735])).
fof(f2735, plain, (spl130_276 <=> (e3 = op(op(e1, e1), op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl130_276])])).
fof(f9450, plain, (~ spl130_83 | ~ spl130_93), inference(avatar_split_clause, [], [f9438, f1587, f1545])).
fof(f1545, plain, (spl130_83 <=> (e2 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl130_83])])).
fof(f9438, plain, (~ (e2 = op(e1, e3)) | ~ spl130_93), inference(backward_demodulation, [], [f659, f1589])).
fof(f9449, plain, (~ spl130_43 | ~ spl130_93), inference(avatar_split_clause, [], [f9436, f1587, f1377])).
fof(f9436, plain, (~ (e2 = op(e3, e1)) | ~ spl130_93), inference(backward_demodulation, [], [f609, f1589])).
fof(f609, plain, ~ (op(e1, e1) = op(e3, e1)), inference(cnf_transformation, [], [f4])).
fof(f9415, plain, (~ spl130_45 | ~ spl130_70), inference(avatar_split_clause, [], [f9414, f1490, f1385])).
fof(f9414, plain, (~ (e4 = op(e3, e1)) | ~ spl130_70), inference(forward_demodulation, [], [f611, f1492])).
fof(f611, plain, ~ (op(e2, e1) = op(e3, e1)), inference(cnf_transformation, [], [f4])).
fof(f9413, plain, (~ spl130_14 | ~ spl130_64), inference(avatar_split_clause, [], [f9412, f1465, f1255])).
fof(f1255, plain, (spl130_14 <=> (e3 = op(e4, e2))), introduced(avatar_definition, [new_symbols(naming, [spl130_14])])).
fof(f9412, plain, (~ (e3 = op(e4, e2)) | ~ spl130_64), inference(forward_demodulation, [], [f622, f1467])).
fof(f622, plain, ~ (op(e2, e2) = op(e4, e2)), inference(cnf_transformation, [], [f4])).
fof(f9411, plain, (~ spl130_10 | ~ spl130_25), inference(avatar_split_clause, [], [f9410, f1301, f1238])).
fof(f1238, plain, (spl130_10 <=> (e4 = op(e4, e3))), introduced(avatar_definition, [new_symbols(naming, [spl130_10])])).
fof(f9410, plain, (~ (e4 = op(e4, e3)) | ~ spl130_25), inference(forward_demodulation, [], [f686, f1303])).
fof(f686, plain, ~ (op(e4, e0) = op(e4, e3)), inference(cnf_transformation, [], [f4])).
fof(f9408, plain, (~ spl130_9 | ~ spl130_109), inference(avatar_split_clause, [], [f9407, f1654, f1234])).
fof(f1234, plain, (spl130_9 <=> (e3 = op(e4, e3))), introduced(avatar_definition, [new_symbols(naming, [spl130_9])])).
fof(f9407, plain, (~ (e3 = op(e4, e3)) | ~ spl130_109), inference(forward_demodulation, [], [f627, f1656])).
fof(f9375, plain, (~ spl130_31 | ~ spl130_56), inference(avatar_split_clause, [], [f9369, f1432, f1327])).
fof(f1327, plain, (spl130_31 <=> (e0 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl130_31])])).
fof(f9369, plain, (~ (e0 = op(e3, e3)) | ~ spl130_56), inference(backward_demodulation, [], [f631, f1434])).
fof(f1434, plain, ((e0 = op(e2, e3)) | ~ spl130_56), inference(avatar_component_clause, [], [f1432])).
fof(f631, plain, ~ (op(e2, e3) = op(e3, e3)), inference(cnf_transformation, [], [f4])).
fof(f9366, plain, (~ spl130_54 | ~ spl130_64), inference(avatar_split_clause, [], [f9361, f1465, f1423])).
fof(f9361, plain, (~ (e3 = op(e2, e4)) | ~ spl130_64), inference(backward_demodulation, [], [f672, f1467])).
fof(f9358, plain, (~ spl130_60 | ~ spl130_70), inference(avatar_split_clause, [], [f9356, f1490, f1448])).
fof(f9356, plain, (~ (e4 = op(e2, e3)) | ~ spl130_70), inference(backward_demodulation, [], [f669, f1492])).
fof(f669, plain, ~ (op(e2, e1) = op(e2, e3)), inference(cnf_transformation, [], [f4])).
fof(f9350, plain, (~ spl130_40 | ~ spl130_90), inference(avatar_split_clause, [], [f9348, f1574, f1364])).
fof(f1364, plain, (spl130_40 <=> (e4 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl130_40])])).
fof(f9348, plain, (~ (e4 = op(e3, e2)) | ~ spl130_90), inference(backward_demodulation, [], [f619, f1576])).
fof(f9346, plain, (~ spl130_91 | ~ spl130_121 | spl130_312), inference(avatar_contradiction_clause, [], [f9345])).
fof(f9345, plain, ($false | (~ spl130_91 | ~ spl130_121 | spl130_312)), inference(subsumption_resolution, [], [f9335, f1707])).
fof(f9335, plain, (~ (e0 = op(e0, e0)) | (~ spl130_91 | spl130_312)), inference(backward_demodulation, [], [f2937, f1581])).
fof(f2937, plain, (~ (e0 = op(op(e1, e1), op(e1, e1))) | spl130_312), inference(avatar_component_clause, [], [f2935])).
fof(f2935, plain, (spl130_312 <=> (e0 = op(op(e1, e1), op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl130_312])])).
fof(f9337, plain, (~ spl130_76 | ~ spl130_91), inference(avatar_split_clause, [], [f9322, f1579, f1516])).
fof(f1516, plain, (spl130_76 <=> (e0 = op(e1, e4))), introduced(avatar_definition, [new_symbols(naming, [spl130_76])])).
fof(f9322, plain, (~ (e0 = op(e1, e4)) | ~ spl130_91), inference(backward_demodulation, [], [f660, f1581])).
fof(f9336, plain, (~ spl130_66 | ~ spl130_91), inference(avatar_split_clause, [], [f9320, f1579, f1474])).
fof(f9320, plain, (~ (e0 = op(e2, e1)) | ~ spl130_91), inference(backward_demodulation, [], [f608, f1581])).
fof(f608, plain, ~ (op(e1, e1) = op(e2, e1)), inference(cnf_transformation, [], [f4])).
fof(f9319, plain, (~ spl130_91 | spl130_261 | ~ spl130_262), inference(avatar_split_clause, [], [f9317, f2672, f2668, f1579])).
fof(f2668, plain, (spl130_261 <=> (e0 = op(op(op(e2, e2), e2), op(op(e2, e2), e2)))), introduced(avatar_definition, [new_symbols(naming, [spl130_261])])).
fof(f2672, plain, (spl130_262 <=> (e1 = op(op(e2, e2), e2))), introduced(avatar_definition, [new_symbols(naming, [spl130_262])])).
fof(f9317, plain, (~ (e0 = op(e1, e1)) | (spl130_261 | ~ spl130_262)), inference(backward_demodulation, [], [f2670, f2673])).
fof(f2673, plain, ((e1 = op(op(e2, e2), e2)) | ~ spl130_262), inference(avatar_component_clause, [], [f2672])).
fof(f2670, plain, (~ (e0 = op(op(op(e2, e2), e2), op(op(e2, e2), e2))) | spl130_261), inference(avatar_component_clause, [], [f2668])).
fof(f9313, plain, (~ spl130_91 | spl130_265 | ~ spl130_266), inference(avatar_split_clause, [], [f9308, f2690, f2686, f1579])).
fof(f2686, plain, (spl130_265 <=> (e0 = op(op(op(e3, e3), e3), op(op(e3, e3), e3)))), introduced(avatar_definition, [new_symbols(naming, [spl130_265])])).
fof(f2690, plain, (spl130_266 <=> (e1 = op(op(e3, e3), e3))), introduced(avatar_definition, [new_symbols(naming, [spl130_266])])).
fof(f9308, plain, (~ (e0 = op(e1, e1)) | (spl130_265 | ~ spl130_266)), inference(backward_demodulation, [], [f2688, f2691])).
fof(f2691, plain, ((e1 = op(op(e3, e3), e3)) | ~ spl130_266), inference(avatar_component_clause, [], [f2690])).
fof(f2688, plain, (~ (e0 = op(op(op(e3, e3), e3), op(op(e3, e3), e3))) | spl130_265), inference(avatar_component_clause, [], [f2686])).
fof(f9268, plain, (~ spl130_29 | ~ spl130_49), inference(avatar_split_clause, [], [f9267, f1402, f1318])).
fof(f1318, plain, (spl130_29 <=> (e3 = op(e3, e4))), introduced(avatar_definition, [new_symbols(naming, [spl130_29])])).
fof(f9267, plain, (~ (e3 = op(e3, e4)) | ~ spl130_49), inference(forward_demodulation, [], [f677, f1404])).
fof(f9266, plain, (~ spl130_30 | ~ spl130_105), inference(avatar_split_clause, [], [f9265, f1637, f1322])).
fof(f1322, plain, (spl130_30 <=> (e4 = op(e3, e4))), introduced(avatar_definition, [new_symbols(naming, [spl130_30])])).
fof(f9265, plain, (~ (e4 = op(e3, e4)) | ~ spl130_105), inference(forward_demodulation, [], [f636, f1639])).
fof(f9264, plain, (~ spl130_3 | ~ spl130_62 | spl130_280), inference(avatar_contradiction_clause, [], [f9263])).
fof(f9263, plain, ($false | (~ spl130_3 | ~ spl130_62 | spl130_280)), inference(subsumption_resolution, [], [f9258, f1459])).
fof(f9258, plain, (~ (e1 = op(e2, e2)) | (~ spl130_3 | spl130_280)), inference(backward_demodulation, [], [f2757, f1211])).
fof(f1211, plain, ((e2 = op(e4, e4)) | ~ spl130_3), inference(avatar_component_clause, [], [f1209])).
fof(f2757, plain, (~ (e1 = op(op(e4, e4), op(e4, e4))) | spl130_280), inference(avatar_component_clause, [], [f2755])).
fof(f2755, plain, (spl130_280 <=> (e1 = op(op(e4, e4), op(e4, e4)))), introduced(avatar_definition, [new_symbols(naming, [spl130_280])])).
fof(f9178, plain, (~ spl130_23 | ~ spl130_73), inference(avatar_split_clause, [], [f9171, f1503, f1293])).
fof(f1293, plain, (spl130_23 <=> (e2 = op(e4, e0))), introduced(avatar_definition, [new_symbols(naming, [spl130_23])])).
fof(f9171, plain, (~ (e2 = op(e4, e0)) | ~ spl130_73), inference(backward_demodulation, [], [f602, f1505])).
fof(f9162, plain, (~ spl130_78 | ~ spl130_83), inference(avatar_split_clause, [], [f9158, f1545, f1524])).
fof(f9158, plain, (~ (e2 = op(e1, e4)) | ~ spl130_83), inference(backward_demodulation, [], [f663, f1547])).
fof(f1547, plain, ((e2 = op(e1, e3)) | ~ spl130_83), inference(avatar_component_clause, [], [f1545])).
fof(f663, plain, ~ (op(e1, e3) = op(e1, e4)), inference(cnf_transformation, [], [f4])).
fof(f9161, plain, (~ spl130_33 | ~ spl130_83), inference(avatar_split_clause, [], [f9157, f1545, f1335])).
fof(f9157, plain, (~ (e2 = op(e3, e3)) | ~ spl130_83), inference(backward_demodulation, [], [f629, f1547])).
fof(f629, plain, ~ (op(e1, e3) = op(e3, e3)), inference(cnf_transformation, [], [f4])).
fof(f9155, plain, (~ spl130_39 | ~ spl130_89), inference(avatar_split_clause, [], [f9150, f1570, f1360])).
fof(f1570, plain, (spl130_89 <=> (e3 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl130_89])])).
fof(f9150, plain, (~ (e3 = op(e3, e2)) | ~ spl130_89), inference(backward_demodulation, [], [f619, f1572])).
fof(f1572, plain, ((e3 = op(e1, e2)) | ~ spl130_89), inference(avatar_component_clause, [], [f1570])).
fof(f9045, plain, (~ spl130_31 | spl130_261 | ~ spl130_283), inference(avatar_split_clause, [], [f9044, f2770, f2668, f1327])).
fof(f9044, plain, (~ (e0 = op(e3, e3)) | (spl130_261 | ~ spl130_283)), inference(backward_demodulation, [], [f2670, f2771])).
fof(f9028, plain, (~ spl130_100 | ~ spl130_95), inference(avatar_split_clause, [], [f9027, f1595, f1616])).
fof(f1616, plain, (spl130_100 <=> (e4 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl130_100])])).
fof(f1595, plain, (spl130_95 <=> (e4 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl130_95])])).
fof(f9027, plain, (~ (e4 = op(e1, e0)) | ~ spl130_95), inference(forward_demodulation, [], [f654, f1597])).
fof(f1597, plain, ((e4 = op(e1, e1)) | ~ spl130_95), inference(avatar_component_clause, [], [f1595])).
fof(f9026, plain, (~ spl130_90 | ~ spl130_95), inference(avatar_split_clause, [], [f9025, f1595, f1574])).
fof(f9025, plain, (~ (e4 = op(e1, e2)) | ~ spl130_95), inference(forward_demodulation, [], [f658, f1597])).
fof(f658, plain, ~ (op(e1, e1) = op(e1, e2)), inference(cnf_transformation, [], [f4])).
fof(f9020, plain, (~ spl130_70 | ~ spl130_95), inference(avatar_split_clause, [], [f9019, f1595, f1490])).
fof(f9019, plain, (~ (e4 = op(e2, e1)) | ~ spl130_95), inference(forward_demodulation, [], [f608, f1597])).
fof(f9005, plain, (~ spl130_3 | ~ spl130_95 | spl130_282), inference(avatar_split_clause, [], [f8790, f2765, f1595, f1209])).
fof(f2765, plain, (spl130_282 <=> (e2 = op(op(e1, e1), op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl130_282])])).
fof(f8790, plain, (~ (e2 = op(e4, e4)) | (~ spl130_95 | spl130_282)), inference(backward_demodulation, [], [f2767, f1597])).
fof(f2767, plain, (~ (e2 = op(op(e1, e1), op(e1, e1))) | spl130_282), inference(avatar_component_clause, [], [f2765])).
fof(f8984, plain, (~ spl130_8 | ~ spl130_23), inference(avatar_split_clause, [], [f8980, f1293, f1230])).
fof(f8980, plain, (~ (e2 = op(e4, e3)) | ~ spl130_23), inference(backward_demodulation, [], [f686, f1295])).
fof(f1295, plain, ((e2 = op(e4, e0)) | ~ spl130_23), inference(avatar_component_clause, [], [f1293])).
fof(f8972, plain, (~ spl130_8 | ~ spl130_33), inference(avatar_split_clause, [], [f8964, f1335, f1230])).
fof(f8964, plain, (~ (e2 = op(e4, e3)) | ~ spl130_33), inference(backward_demodulation, [], [f633, f1337])).
fof(f633, plain, ~ (op(e3, e3) = op(e4, e3)), inference(cnf_transformation, [], [f4])).
fof(f8962, plain, (~ spl130_16 | ~ spl130_41), inference(avatar_split_clause, [], [f8955, f1369, f1264])).
fof(f1369, plain, (spl130_41 <=> (e0 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl130_41])])).
fof(f8955, plain, (~ (e0 = op(e4, e1)) | ~ spl130_41), inference(backward_demodulation, [], [f613, f1371])).
fof(f1371, plain, ((e0 = op(e3, e1)) | ~ spl130_41), inference(avatar_component_clause, [], [f1369])).
fof(f613, plain, ~ (op(e3, e1) = op(e4, e1)), inference(cnf_transformation, [], [f4])).
fof(f8950, plain, (~ spl130_43 | ~ spl130_118), inference(avatar_split_clause, [], [f8945, f1692, f1377])).
fof(f1692, plain, (spl130_118 <=> (e2 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl130_118])])).
fof(f8945, plain, (~ (e2 = op(e3, e1)) | ~ spl130_118), inference(backward_demodulation, [], [f606, f1694])).
fof(f1694, plain, ((e2 = op(e0, e1)) | ~ spl130_118), inference(avatar_component_clause, [], [f1692])).
fof(f606, plain, ~ (op(e0, e1) = op(e3, e1)), inference(cnf_transformation, [], [f4])).
fof(f8944, plain, (~ spl130_95 | ~ spl130_122 | spl130_297), inference(avatar_contradiction_clause, [], [f8943])).
fof(f8943, plain, ($false | (~ spl130_95 | ~ spl130_122 | spl130_297)), inference(subsumption_resolution, [], [f8935, f1597])).
fof(f8935, plain, (~ (e4 = op(e1, e1)) | (~ spl130_122 | spl130_297)), inference(backward_demodulation, [], [f2844, f1711])).
fof(f1711, plain, ((op(e0, e0) = e1) | ~ spl130_122), inference(avatar_component_clause, [], [f1709])).
fof(f1709, plain, (spl130_122 <=> (op(e0, e0) = e1)), introduced(avatar_definition, [new_symbols(naming, [spl130_122])])).
fof(f2844, plain, (~ (e4 = op(op(e0, e0), op(e0, e0))) | spl130_297), inference(avatar_component_clause, [], [f2842])).
fof(f2842, plain, (spl130_297 <=> (e4 = op(op(e0, e0), op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl130_297])])).
fof(f8942, plain, (~ spl130_102 | ~ spl130_122), inference(avatar_split_clause, [], [f8931, f1709, f1625])).
fof(f8931, plain, (~ (e1 = op(e0, e4)) | ~ spl130_122), inference(backward_demodulation, [], [f647, f1711])).
fof(f8941, plain, (~ spl130_22 | ~ spl130_122), inference(avatar_split_clause, [], [f8929, f1709, f1289])).
fof(f8929, plain, (~ (e1 = op(e4, e0)) | ~ spl130_122), inference(backward_demodulation, [], [f597, f1711])).
fof(f597, plain, ~ (op(e0, e0) = op(e4, e0)), inference(cnf_transformation, [], [f4])).
fof(f8916, plain, (spl130_33 | ~ spl130_301 | ~ spl130_309), inference(avatar_split_clause, [], [f8914, f2914, f2864, f1335])).
fof(f2864, plain, (spl130_301 <=> (e3 = op(op(e0, e0), e0))), introduced(avatar_definition, [new_symbols(naming, [spl130_301])])).
fof(f8914, plain, ((e2 = op(e3, e3)) | (~ spl130_301 | ~ spl130_309)), inference(backward_demodulation, [], [f2915, f2865])).
fof(f2865, plain, ((e3 = op(op(e0, e0), e0)) | ~ spl130_301), inference(avatar_component_clause, [], [f2864])).
fof(f2915, plain, ((e2 = op(op(op(e0, e0), e0), op(op(e0, e0), e0))) | ~ spl130_309), inference(avatar_component_clause, [], [f2914])).
fof(f8894, plain, (~ spl130_10 | ~ spl130_110), inference(avatar_split_clause, [], [f8893, f1658, f1238])).
fof(f1658, plain, (spl130_110 <=> (e4 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl130_110])])).
fof(f8893, plain, (~ (e4 = op(e4, e3)) | ~ spl130_110), inference(forward_demodulation, [], [f627, f1660])).
fof(f1660, plain, ((e4 = op(e0, e3)) | ~ spl130_110), inference(avatar_component_clause, [], [f1658])).
fof(f8891, plain, (~ spl130_123 | ~ spl130_294 | spl130_308), inference(avatar_split_clause, [], [f8887, f2908, f2828, f1713])).
fof(f2828, plain, (spl130_294 <=> (e0 = op(op(e4, e4), e4))), introduced(avatar_definition, [new_symbols(naming, [spl130_294])])).
fof(f2908, plain, (spl130_308 <=> (e2 = op(op(op(e4, e4), e4), op(op(e4, e4), e4)))), introduced(avatar_definition, [new_symbols(naming, [spl130_308])])).
fof(f8887, plain, (~ (op(e0, e0) = e2) | (~ spl130_294 | spl130_308)), inference(backward_demodulation, [], [f2910, f2829])).
fof(f2829, plain, ((e0 = op(op(e4, e4), e4)) | ~ spl130_294), inference(avatar_component_clause, [], [f2828])).
fof(f2910, plain, (~ (e2 = op(op(op(e4, e4), e4), op(op(e4, e4), e4))) | spl130_308), inference(avatar_component_clause, [], [f2908])).
fof(f8889, plain, (~ spl130_124 | ~ spl130_294 | spl130_315), inference(avatar_split_clause, [], [f8885, f2960, f2828, f1717])).
fof(f2960, plain, (spl130_315 <=> (e3 = op(op(op(e4, e4), e4), op(op(e4, e4), e4)))), introduced(avatar_definition, [new_symbols(naming, [spl130_315])])).
fof(f8885, plain, (~ (op(e0, e0) = e3) | (~ spl130_294 | spl130_315)), inference(backward_demodulation, [], [f2962, f2829])).
fof(f2962, plain, (~ (e3 = op(op(op(e4, e4), e4), op(op(e4, e4), e4))) | spl130_315), inference(avatar_component_clause, [], [f2960])).
fof(f8873, plain, (~ spl130_16 | ~ spl130_95 | ~ spl130_123 | spl130_305), inference(avatar_contradiction_clause, [], [f8872])).
fof(f8872, plain, ($false | (~ spl130_16 | ~ spl130_95 | ~ spl130_123 | spl130_305)), inference(subsumption_resolution, [], [f8871, f1715])).
fof(f8871, plain, (~ (op(e0, e0) = e2) | (~ spl130_16 | ~ spl130_95 | spl130_305)), inference(forward_demodulation, [], [f8870, f1266])).
fof(f8870, plain, (~ (e2 = op(op(e4, e1), op(e4, e1))) | (~ spl130_95 | spl130_305)), inference(forward_demodulation, [], [f2894, f1597])).
fof(f2894, plain, (~ (e2 = op(op(op(e1, e1), e1), op(op(e1, e1), e1))) | spl130_305), inference(avatar_component_clause, [], [f2892])).
fof(f2892, plain, (spl130_305 <=> (e2 = op(op(op(e1, e1), e1), op(op(e1, e1), e1)))), introduced(avatar_definition, [new_symbols(naming, [spl130_305])])).
fof(f8805, plain, (~ spl130_76 | ~ spl130_81), inference(avatar_split_clause, [], [f8799, f1537, f1516])).
fof(f8799, plain, (~ (e0 = op(e1, e4)) | ~ spl130_81), inference(backward_demodulation, [], [f663, f1539])).
fof(f8804, plain, (~ spl130_31 | ~ spl130_81), inference(avatar_split_clause, [], [f8798, f1537, f1327])).
fof(f8798, plain, (~ (e0 = op(e3, e3)) | ~ spl130_81), inference(backward_demodulation, [], [f629, f1539])).
fof(f8797, plain, (~ spl130_4 | ~ spl130_95 | spl130_276), inference(avatar_contradiction_clause, [], [f8796])).
fof(f8796, plain, ($false | (~ spl130_4 | ~ spl130_95 | spl130_276)), inference(subsumption_resolution, [], [f8789, f1215])).
fof(f1215, plain, ((e3 = op(e4, e4)) | ~ spl130_4), inference(avatar_component_clause, [], [f1213])).
fof(f1213, plain, (spl130_4 <=> (e3 = op(e4, e4))), introduced(avatar_definition, [new_symbols(naming, [spl130_4])])).
fof(f8789, plain, (~ (e3 = op(e4, e4)) | (~ spl130_95 | spl130_276)), inference(backward_demodulation, [], [f2737, f1597])).
fof(f8794, plain, (~ spl130_85 | ~ spl130_95), inference(avatar_split_clause, [], [f8786, f1595, f1553])).
fof(f8786, plain, (~ (e4 = op(e1, e3)) | ~ spl130_95), inference(backward_demodulation, [], [f659, f1597])).
fof(f8784, plain, (~ spl130_27 | ~ spl130_102), inference(avatar_split_clause, [], [f8780, f1625, f1310])).
fof(f8780, plain, (~ (e1 = op(e3, e4)) | ~ spl130_102), inference(backward_demodulation, [], [f636, f1627])).
fof(f1627, plain, ((e1 = op(e0, e4)) | ~ spl130_102), inference(avatar_component_clause, [], [f1625])).
fof(f8778, plain, (~ spl130_85 | ~ spl130_110), inference(avatar_split_clause, [], [f8775, f1658, f1553])).
fof(f8775, plain, (~ (e4 = op(e1, e3)) | ~ spl130_110), inference(backward_demodulation, [], [f624, f1660])).
fof(f8732, plain, (~ spl130_94 | ~ spl130_270 | spl130_315), inference(avatar_split_clause, [], [f8575, f2960, f2708, f1591])).
fof(f2708, plain, (spl130_270 <=> (e1 = op(op(e4, e4), e4))), introduced(avatar_definition, [new_symbols(naming, [spl130_270])])).
fof(f8575, plain, (~ (e3 = op(e1, e1)) | (~ spl130_270 | spl130_315)), inference(forward_demodulation, [], [f2962, f2709])).
fof(f2709, plain, ((e1 = op(op(e4, e4), e4)) | ~ spl130_270), inference(avatar_component_clause, [], [f2708])).
fof(f8731, plain, (~ spl130_94 | ~ spl130_99), inference(avatar_split_clause, [], [f8730, f1612, f1591])).
fof(f1612, plain, (spl130_99 <=> (e3 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl130_99])])).
fof(f8730, plain, (~ (e3 = op(e1, e1)) | ~ spl130_99), inference(forward_demodulation, [], [f654, f1614])).
fof(f1614, plain, ((e3 = op(e1, e0)) | ~ spl130_99), inference(avatar_component_clause, [], [f1612])).
fof(f8683, plain, (~ spl130_35 | ~ spl130_45), inference(avatar_split_clause, [], [f8681, f1385, f1343])).
fof(f8681, plain, (~ (e4 = op(e3, e3)) | ~ spl130_45), inference(backward_demodulation, [], [f679, f1387])).
fof(f1387, plain, ((e4 = op(e3, e1)) | ~ spl130_45), inference(avatar_component_clause, [], [f1385])).
fof(f679, plain, ~ (op(e3, e1) = op(e3, e3)), inference(cnf_transformation, [], [f4])).
fof(f8678, plain, (~ spl130_23 | ~ spl130_48), inference(avatar_split_clause, [], [f8672, f1398, f1293])).
fof(f8672, plain, (~ (e2 = op(e4, e0)) | ~ spl130_48), inference(backward_demodulation, [], [f603, f1400])).
fof(f8660, plain, (~ spl130_79 | ~ spl130_99), inference(avatar_split_clause, [], [f8658, f1612, f1528])).
fof(f8658, plain, (~ (e3 = op(e1, e4)) | ~ spl130_99), inference(backward_demodulation, [], [f657, f1614])).
fof(f8648, plain, (~ spl130_102 | ~ spl130_107), inference(avatar_split_clause, [], [f8643, f1646, f1625])).
fof(f8643, plain, (~ (e1 = op(e0, e4)) | ~ spl130_107), inference(backward_demodulation, [], [f653, f1648])).
fof(f653, plain, ~ (op(e0, e3) = op(e0, e4)), inference(cnf_transformation, [], [f4])).
fof(f8639, plain, (~ spl130_19 | ~ spl130_119), inference(avatar_split_clause, [], [f8635, f1696, f1276])).
fof(f1276, plain, (spl130_19 <=> (e3 = op(e4, e1))), introduced(avatar_definition, [new_symbols(naming, [spl130_19])])).
fof(f1696, plain, (spl130_119 <=> (e3 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl130_119])])).
fof(f8635, plain, (~ (e3 = op(e4, e1)) | ~ spl130_119), inference(backward_demodulation, [], [f607, f1698])).
fof(f1698, plain, ((e3 = op(e0, e1)) | ~ spl130_119), inference(avatar_component_clause, [], [f1696])).
fof(f607, plain, ~ (op(e0, e1) = op(e4, e1)), inference(cnf_transformation, [], [f4])).
fof(f8632, plain, (~ spl130_4 | ~ spl130_125 | spl130_298), inference(avatar_split_clause, [], [f8624, f2847, f1721, f1213])).
fof(f1721, plain, (spl130_125 <=> (op(e0, e0) = e4)), introduced(avatar_definition, [new_symbols(naming, [spl130_125])])).
fof(f2847, plain, (spl130_298 <=> (e3 = op(op(e0, e0), op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl130_298])])).
fof(f8624, plain, (~ (e3 = op(e4, e4)) | (~ spl130_125 | spl130_298)), inference(backward_demodulation, [], [f2849, f1723])).
fof(f1723, plain, ((op(e0, e0) = e4) | ~ spl130_125), inference(avatar_component_clause, [], [f1721])).
fof(f2849, plain, (~ (e3 = op(op(e0, e0), op(e0, e0))) | spl130_298), inference(avatar_component_clause, [], [f2847])).
fof(f8629, plain, (~ spl130_120 | ~ spl130_125), inference(avatar_split_clause, [], [f8620, f1721, f1700])).
fof(f8620, plain, (~ (e4 = op(e0, e1)) | ~ spl130_125), inference(backward_demodulation, [], [f644, f1723])).
fof(f8628, plain, (~ spl130_100 | ~ spl130_125), inference(avatar_split_clause, [], [f8617, f1721, f1616])).
fof(f8617, plain, (~ (e4 = op(e1, e0)) | ~ spl130_125), inference(backward_demodulation, [], [f594, f1723])).
fof(f594, plain, ~ (op(e0, e0) = op(e1, e0)), inference(cnf_transformation, [], [f4])).
fof(f8590, plain, (~ spl130_42 | ~ spl130_67), inference(avatar_split_clause, [], [f8589, f1478, f1373])).
fof(f1373, plain, (spl130_42 <=> (e1 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl130_42])])).
fof(f8589, plain, (~ (e1 = op(e3, e1)) | ~ spl130_67), inference(forward_demodulation, [], [f611, f1480])).
fof(f1480, plain, ((e1 = op(e2, e1)) | ~ spl130_67), inference(avatar_component_clause, [], [f1478])).
fof(f8588, plain, (~ spl130_44 | ~ spl130_39), inference(avatar_split_clause, [], [f8530, f1360, f1381])).
fof(f8530, plain, (~ (e3 = op(e3, e1)) | ~ spl130_39), inference(forward_demodulation, [], [f678, f1362])).
fof(f1362, plain, ((e3 = op(e3, e2)) | ~ spl130_39), inference(avatar_component_clause, [], [f1360])).
fof(f678, plain, ~ (op(e3, e1) = op(e3, e2)), inference(cnf_transformation, [], [f4])).
fof(f8554, plain, (~ spl130_47 | ~ spl130_124 | spl130_310), inference(avatar_contradiction_clause, [], [f8553])).
fof(f8553, plain, ($false | (~ spl130_47 | ~ spl130_124 | spl130_310)), inference(subsumption_resolution, [], [f8552, f1396])).
fof(f1396, plain, ((e1 = op(e3, e0)) | ~ spl130_47), inference(avatar_component_clause, [], [f1394])).
fof(f1394, plain, (spl130_47 <=> (e1 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl130_47])])).
fof(f8552, plain, (~ (e1 = op(e3, e0)) | (~ spl130_124 | spl130_310)), inference(forward_demodulation, [], [f2920, f1719])).
fof(f2920, plain, (~ (e1 = op(op(e0, e0), e0)) | spl130_310), inference(avatar_component_clause, [], [f2918])).
fof(f2918, plain, (spl130_310 <=> (e1 = op(op(e0, e0), e0))), introduced(avatar_definition, [new_symbols(naming, [spl130_310])])).
fof(f8514, plain, (~ spl130_9 | ~ spl130_19), inference(avatar_split_clause, [], [f8512, f1276, f1234])).
fof(f8512, plain, (~ (e3 = op(e4, e3)) | ~ spl130_19), inference(backward_demodulation, [], [f689, f1278])).
fof(f1278, plain, ((e3 = op(e4, e1)) | ~ spl130_19), inference(avatar_component_clause, [], [f1276])).
fof(f689, plain, ~ (op(e4, e1) = op(e4, e3)), inference(cnf_transformation, [], [f4])).
fof(f8510, plain, (~ spl130_13 | ~ spl130_23), inference(avatar_split_clause, [], [f8506, f1293, f1251])).
fof(f1251, plain, (spl130_13 <=> (e2 = op(e4, e2))), introduced(avatar_definition, [new_symbols(naming, [spl130_13])])).
fof(f8506, plain, (~ (e2 = op(e4, e2)) | ~ spl130_23), inference(backward_demodulation, [], [f685, f1295])).
fof(f8488, plain, (~ spl130_29 | ~ spl130_39), inference(avatar_split_clause, [], [f8486, f1360, f1318])).
fof(f8486, plain, (~ (e3 = op(e3, e4)) | ~ spl130_39), inference(backward_demodulation, [], [f682, f1362])).
fof(f8483, plain, (~ spl130_41 | ~ spl130_43), inference(avatar_contradiction_clause, [], [f8482])).
fof(f8482, plain, ($false | (~ spl130_41 | ~ spl130_43)), inference(subsumption_resolution, [], [f8481, f695])).
fof(f695, plain, ~ (e0 = e2), inference(cnf_transformation, [], [f5])).
fof(f5, plain, (~ (e3 = e4) & ~ (e2 = e4) & ~ (e2 = e3) & ~ (e1 = e4) & ~ (e1 = e3) & ~ (e1 = e2) & ~ (e0 = e4) & ~ (e0 = e3) & ~ (e0 = e2) & ~ (e0 = e1)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG099+1.p', ax5)).
fof(f8481, plain, ((e0 = e2) | (~ spl130_41 | ~ spl130_43)), inference(backward_demodulation, [], [f1379, f1371])).
fof(f8479, plain, (~ spl130_37 | ~ spl130_47), inference(avatar_split_clause, [], [f8473, f1394, f1352])).
fof(f8473, plain, (~ (e1 = op(e3, e2)) | ~ spl130_47), inference(backward_demodulation, [], [f675, f1396])).
fof(f8470, plain, (~ spl130_9 | ~ spl130_59), inference(avatar_split_clause, [], [f8467, f1444, f1234])).
fof(f8467, plain, (~ (e3 = op(e4, e3)) | ~ spl130_59), inference(backward_demodulation, [], [f632, f1446])).
fof(f1446, plain, ((e3 = op(e2, e3)) | ~ spl130_59), inference(avatar_component_clause, [], [f1444])).
fof(f8463, plain, (~ spl130_53 | ~ spl130_63), inference(avatar_split_clause, [], [f8457, f1461, f1419])).
fof(f1419, plain, (spl130_53 <=> (e2 = op(e2, e4))), introduced(avatar_definition, [new_symbols(naming, [spl130_53])])).
fof(f8457, plain, (~ (e2 = op(e2, e4)) | ~ spl130_63), inference(backward_demodulation, [], [f672, f1463])).
fof(f1463, plain, ((e2 = op(e2, e2)) | ~ spl130_63), inference(avatar_component_clause, [], [f1461])).
fof(f8462, plain, (~ spl130_13 | ~ spl130_63), inference(avatar_split_clause, [], [f8455, f1461, f1251])).
fof(f8455, plain, (~ (e2 = op(e4, e2)) | ~ spl130_63), inference(backward_demodulation, [], [f622, f1463])).
fof(f8453, plain, (~ spl130_17 | ~ spl130_67), inference(avatar_split_clause, [], [f8445, f1478, f1268])).
fof(f1268, plain, (spl130_17 <=> (e1 = op(e4, e1))), introduced(avatar_definition, [new_symbols(naming, [spl130_17])])).
fof(f8445, plain, (~ (e1 = op(e4, e1)) | ~ spl130_67), inference(backward_demodulation, [], [f612, f1480])).
fof(f8443, plain, (~ spl130_21 | ~ spl130_71), inference(avatar_split_clause, [], [f8434, f1495, f1285])).
fof(f1285, plain, (spl130_21 <=> (e0 = op(e4, e0))), introduced(avatar_definition, [new_symbols(naming, [spl130_21])])).
fof(f8434, plain, (~ (e0 = op(e4, e0)) | ~ spl130_71), inference(backward_demodulation, [], [f602, f1497])).
fof(f1497, plain, ((e0 = op(e2, e0)) | ~ spl130_71), inference(avatar_component_clause, [], [f1495])).
fof(f8432, plain, (~ spl130_29 | ~ spl130_79), inference(avatar_split_clause, [], [f8430, f1528, f1318])).
fof(f8430, plain, (~ (e3 = op(e3, e4)) | ~ spl130_79), inference(backward_demodulation, [], [f639, f1530])).
fof(f1530, plain, ((e3 = op(e1, e4)) | ~ spl130_79), inference(avatar_component_clause, [], [f1528])).
fof(f8419, plain, (~ spl130_77 | ~ spl130_87), inference(avatar_split_clause, [], [f8414, f1562, f1520])).
fof(f1520, plain, (spl130_77 <=> (e1 = op(e1, e4))), introduced(avatar_definition, [new_symbols(naming, [spl130_77])])).
fof(f8414, plain, (~ (e1 = op(e1, e4)) | ~ spl130_87), inference(backward_demodulation, [], [f662, f1564])).
fof(f8409, plain, (~ spl130_66 | ~ spl130_93 | spl130_306), inference(avatar_split_clause, [], [f8405, f2896, f1587, f1474])).
fof(f2896, plain, (spl130_306 <=> (e0 = op(op(e1, e1), e1))), introduced(avatar_definition, [new_symbols(naming, [spl130_306])])).
fof(f8405, plain, (~ (e0 = op(e2, e1)) | (~ spl130_93 | spl130_306)), inference(backward_demodulation, [], [f2898, f1589])).
fof(f2898, plain, (~ (e0 = op(op(e1, e1), e1)) | spl130_306), inference(avatar_component_clause, [], [f2896])).
fof(f8391, plain, (~ spl130_50 | ~ spl130_100), inference(avatar_split_clause, [], [f8385, f1616, f1406])).
fof(f8385, plain, (~ (e4 = op(e3, e0)) | ~ spl130_100), inference(backward_demodulation, [], [f599, f1618])).
fof(f1618, plain, ((e4 = op(e1, e0)) | ~ spl130_100), inference(avatar_component_clause, [], [f1616])).
fof(f599, plain, ~ (op(e1, e0) = op(e3, e0)), inference(cnf_transformation, [], [f4])).
fof(f8374, plain, (~ spl130_83 | ~ spl130_108), inference(avatar_split_clause, [], [f8367, f1650, f1545])).
fof(f8367, plain, (~ (e2 = op(e1, e3)) | ~ spl130_108), inference(backward_demodulation, [], [f624, f1652])).
fof(f8366, plain, (~ spl130_101 | ~ spl130_111), inference(avatar_split_clause, [], [f8360, f1663, f1621])).
fof(f1621, plain, (spl130_101 <=> (e0 = op(e0, e4))), introduced(avatar_definition, [new_symbols(naming, [spl130_101])])).
fof(f8360, plain, (~ (e0 = op(e0, e4)) | ~ spl130_111), inference(backward_demodulation, [], [f652, f1665])).
fof(f1665, plain, ((e0 = op(e0, e2)) | ~ spl130_111), inference(avatar_component_clause, [], [f1663])).
fof(f8354, plain, (~ spl130_115 | ~ spl130_120), inference(avatar_split_clause, [], [f8350, f1700, f1679])).
fof(f8350, plain, (~ (e4 = op(e0, e2)) | ~ spl130_120), inference(backward_demodulation, [], [f648, f1702])).
fof(f8346, plain, (spl130_47 | ~ spl130_124 | ~ spl130_310), inference(avatar_split_clause, [], [f8341, f2918, f1717, f1394])).
fof(f8341, plain, ((e1 = op(e3, e0)) | (~ spl130_124 | ~ spl130_310)), inference(backward_demodulation, [], [f2919, f1719])).
fof(f2919, plain, ((e1 = op(op(e0, e0), e0)) | ~ spl130_310), inference(avatar_component_clause, [], [f2918])).
fof(f8345, plain, (~ spl130_35 | ~ spl130_124 | spl130_297), inference(avatar_split_clause, [], [f8339, f2842, f1717, f1343])).
fof(f8339, plain, (~ (e4 = op(e3, e3)) | (~ spl130_124 | spl130_297)), inference(backward_demodulation, [], [f2844, f1719])).
fof(f8344, plain, (~ spl130_119 | ~ spl130_124), inference(avatar_split_clause, [], [f8334, f1717, f1696])).
fof(f8334, plain, (~ (e3 = op(e0, e1)) | ~ spl130_124), inference(backward_demodulation, [], [f644, f1719])).
fof(f8343, plain, (~ spl130_99 | ~ spl130_124), inference(avatar_split_clause, [], [f8330, f1717, f1612])).
fof(f8330, plain, (~ (e3 = op(e1, e0)) | ~ spl130_124), inference(backward_demodulation, [], [f594, f1719])).
fof(f8304, plain, (~ spl130_95 | ~ spl130_310 | spl130_320), inference(avatar_split_clause, [], [f8303, f3006, f2918, f1595])).
fof(f8303, plain, (~ (e4 = op(e1, e1)) | (~ spl130_310 | spl130_320)), inference(forward_demodulation, [], [f3008, f2919])).
fof(f8297, plain, (~ spl130_93 | spl130_309 | ~ spl130_310), inference(avatar_split_clause, [], [f7851, f2918, f2914, f1587])).
fof(f7851, plain, (~ (e2 = op(e1, e1)) | (spl130_309 | ~ spl130_310)), inference(backward_demodulation, [], [f2916, f2919])).
fof(f8257, plain, (~ spl130_12 | ~ spl130_37), inference(avatar_split_clause, [], [f8251, f1352, f1247])).
fof(f8251, plain, (~ (e1 = op(e4, e2)) | ~ spl130_37), inference(backward_demodulation, [], [f623, f1354])).
fof(f623, plain, ~ (op(e3, e2) = op(e4, e2)), inference(cnf_transformation, [], [f4])).
fof(f8249, plain, (~ spl130_18 | ~ spl130_43), inference(avatar_split_clause, [], [f8243, f1377, f1272])).
fof(f1272, plain, (spl130_18 <=> (e2 = op(e4, e1))), introduced(avatar_definition, [new_symbols(naming, [spl130_18])])).
fof(f8243, plain, (~ (e2 = op(e4, e1)) | ~ spl130_43), inference(backward_demodulation, [], [f613, f1379])).
fof(f8242, plain, (~ spl130_30 | ~ spl130_50), inference(avatar_split_clause, [], [f8241, f1406, f1322])).
fof(f8241, plain, (~ (e4 = op(e3, e4)) | ~ spl130_50), inference(backward_demodulation, [], [f677, f1408])).
fof(f8233, plain, (~ spl130_10 | ~ spl130_60), inference(avatar_split_clause, [], [f8231, f1448, f1238])).
fof(f8231, plain, (~ (e4 = op(e4, e3)) | ~ spl130_60), inference(backward_demodulation, [], [f632, f1450])).
fof(f8214, plain, (~ spl130_58 | ~ spl130_83), inference(avatar_split_clause, [], [f8208, f1545, f1440])).
fof(f8208, plain, (~ (e2 = op(e2, e3)) | ~ spl130_83), inference(backward_demodulation, [], [f628, f1547])).
fof(f8204, plain, (~ spl130_18 | ~ spl130_95 | spl130_274), inference(avatar_split_clause, [], [f8197, f2726, f1595, f1272])).
fof(f8197, plain, (~ (e2 = op(e4, e1)) | (~ spl130_95 | spl130_274)), inference(backward_demodulation, [], [f2728, f1597])).
fof(f8190, plain, (~ spl130_24 | ~ spl130_99), inference(avatar_split_clause, [], [f8185, f1612, f1297])).
fof(f1297, plain, (spl130_24 <=> (e3 = op(e4, e0))), introduced(avatar_definition, [new_symbols(naming, [spl130_24])])).
fof(f8185, plain, (~ (e3 = op(e4, e0)) | ~ spl130_99), inference(backward_demodulation, [], [f600, f1614])).
fof(f600, plain, ~ (op(e1, e0) = op(e4, e0)), inference(cnf_transformation, [], [f4])).
fof(f8172, plain, (~ spl130_82 | ~ spl130_107), inference(avatar_split_clause, [], [f8164, f1646, f1541])).
fof(f1541, plain, (spl130_82 <=> (e1 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl130_82])])).
fof(f8164, plain, (~ (e1 = op(e1, e3)) | ~ spl130_107), inference(backward_demodulation, [], [f624, f1648])).
fof(f8163, plain, (~ spl130_65 | ~ spl130_115), inference(avatar_split_clause, [], [f8158, f1679, f1469])).
fof(f8158, plain, (~ (e4 = op(e2, e2)) | ~ spl130_115), inference(backward_demodulation, [], [f615, f1681])).
fof(f8157, plain, (~ spl130_114 | ~ spl130_119), inference(avatar_split_clause, [], [f8152, f1696, f1675])).
fof(f1675, plain, (spl130_114 <=> (e3 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl130_114])])).
fof(f8152, plain, (~ (e3 = op(e0, e2)) | ~ spl130_119), inference(backward_demodulation, [], [f648, f1698])).
fof(f8142, plain, (spl130_29 | ~ spl130_130), inference(avatar_split_clause, [], [f8132, f1742, f1318])).
fof(f1742, plain, (spl130_130 <=> (e4 = unit)), introduced(avatar_definition, [new_symbols(naming, [spl130_130])])).
fof(f8132, plain, ((e3 = op(e3, e4)) | ~ spl130_130), inference(backward_demodulation, [], [f540, f1744])).
fof(f1744, plain, ((e4 = unit) | ~ spl130_130), inference(avatar_component_clause, [], [f1742])).
fof(f8141, plain, (spl130_9 | ~ spl130_130), inference(avatar_split_clause, [], [f8131, f1742, f1234])).
fof(f8131, plain, ((e3 = op(e4, e3)) | ~ spl130_130), inference(backward_demodulation, [], [f539, f1744])).
fof(f8140, plain, (spl130_53 | ~ spl130_130), inference(avatar_split_clause, [], [f8130, f1742, f1419])).
fof(f8130, plain, ((e2 = op(e2, e4)) | ~ spl130_130), inference(backward_demodulation, [], [f538, f1744])).
fof(f8139, plain, (spl130_13 | ~ spl130_130), inference(avatar_split_clause, [], [f8129, f1742, f1251])).
fof(f8129, plain, ((e2 = op(e4, e2)) | ~ spl130_130), inference(backward_demodulation, [], [f537, f1744])).
fof(f8138, plain, (spl130_77 | ~ spl130_130), inference(avatar_split_clause, [], [f8128, f1742, f1520])).
fof(f8128, plain, ((e1 = op(e1, e4)) | ~ spl130_130), inference(backward_demodulation, [], [f536, f1744])).
fof(f8137, plain, (spl130_17 | ~ spl130_130), inference(avatar_split_clause, [], [f8127, f1742, f1268])).
fof(f8127, plain, ((e1 = op(e4, e1)) | ~ spl130_130), inference(backward_demodulation, [], [f535, f1744])).
fof(f8136, plain, (spl130_101 | ~ spl130_130), inference(avatar_split_clause, [], [f8126, f1742, f1621])).
fof(f8126, plain, ((e0 = op(e0, e4)) | ~ spl130_130), inference(backward_demodulation, [], [f534, f1744])).
fof(f8135, plain, (spl130_21 | ~ spl130_130), inference(avatar_split_clause, [], [f8125, f1742, f1285])).
fof(f8125, plain, ((e0 = op(e4, e0)) | ~ spl130_130), inference(backward_demodulation, [], [f533, f1744])).
fof(f8111, plain, (~ spl130_65 | ~ spl130_123 | spl130_297), inference(avatar_split_clause, [], [f8110, f2842, f1713, f1469])).
fof(f8110, plain, (~ (e4 = op(e2, e2)) | (~ spl130_123 | spl130_297)), inference(forward_demodulation, [], [f2844, f1715])).
fof(f8103, plain, (~ spl130_104 | ~ spl130_1 | spl130_285), inference(avatar_split_clause, [], [f8078, f2780, f1201, f1633])).
fof(f2780, plain, (spl130_285 <=> (e3 = op(op(e4, e4), e4))), introduced(avatar_definition, [new_symbols(naming, [spl130_285])])).
fof(f8078, plain, (~ (e3 = op(e0, e4)) | (~ spl130_1 | spl130_285)), inference(backward_demodulation, [], [f2782, f1203])).
fof(f2782, plain, (~ (e3 = op(op(e4, e4), e4)) | spl130_285), inference(avatar_component_clause, [], [f2780])).
fof(f8098, plain, (~ spl130_94 | ~ spl130_310 | spl130_316), inference(avatar_split_clause, [], [f7853, f2966, f2918, f1591])).
fof(f7853, plain, (~ (e3 = op(e1, e1)) | (~ spl130_310 | spl130_316)), inference(backward_demodulation, [], [f2968, f2919])).
fof(f8095, plain, (~ spl130_64 | ~ spl130_123 | spl130_298), inference(avatar_split_clause, [], [f7892, f2847, f1713, f1465])).
fof(f7892, plain, (~ (e3 = op(e2, e2)) | (~ spl130_123 | spl130_298)), inference(backward_demodulation, [], [f2849, f1715])).
fof(f8068, plain, (~ spl130_5 | ~ spl130_10), inference(avatar_split_clause, [], [f8067, f1238, f1217])).
fof(f1217, plain, (spl130_5 <=> (e4 = op(e4, e4))), introduced(avatar_definition, [new_symbols(naming, [spl130_5])])).
fof(f8067, plain, (~ (e4 = op(e4, e4)) | ~ spl130_10), inference(backward_demodulation, [], [f693, f1240])).
fof(f1240, plain, ((e4 = op(e4, e3)) | ~ spl130_10), inference(avatar_component_clause, [], [f1238])).
fof(f8054, plain, (~ spl130_13 | ~ spl130_18), inference(avatar_split_clause, [], [f8049, f1272, f1251])).
fof(f8049, plain, (~ (e2 = op(e4, e2)) | ~ spl130_18), inference(backward_demodulation, [], [f688, f1274])).
fof(f1274, plain, ((e2 = op(e4, e1)) | ~ spl130_18), inference(avatar_component_clause, [], [f1272])).
fof(f688, plain, ~ (op(e4, e1) = op(e4, e2)), inference(cnf_transformation, [], [f4])).
fof(f8048, plain, (~ spl130_9 | ~ spl130_24), inference(avatar_split_clause, [], [f8045, f1297, f1234])).
fof(f8045, plain, (~ (e3 = op(e4, e3)) | ~ spl130_24), inference(backward_demodulation, [], [f686, f1299])).
fof(f1299, plain, ((e3 = op(e4, e0)) | ~ spl130_24), inference(avatar_component_clause, [], [f1297])).
fof(f8035, plain, (~ spl130_29 | ~ spl130_34), inference(avatar_split_clause, [], [f8025, f1339, f1318])).
fof(f1339, plain, (spl130_34 <=> (e3 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl130_34])])).
fof(f8025, plain, (~ (e3 = op(e3, e4)) | ~ spl130_34), inference(backward_demodulation, [], [f683, f1341])).
fof(f1341, plain, ((e3 = op(e3, e3)) | ~ spl130_34), inference(avatar_component_clause, [], [f1339])).
fof(f8022, plain, (~ spl130_13 | ~ spl130_38), inference(avatar_split_clause, [], [f8017, f1356, f1251])).
fof(f1356, plain, (spl130_38 <=> (e2 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl130_38])])).
fof(f8017, plain, (~ (e2 = op(e4, e2)) | ~ spl130_38), inference(backward_demodulation, [], [f623, f1358])).
fof(f1358, plain, ((e2 = op(e3, e2)) | ~ spl130_38), inference(avatar_component_clause, [], [f1356])).
fof(f8013, plain, (~ spl130_17 | ~ spl130_42), inference(avatar_split_clause, [], [f8003, f1373, f1268])).
fof(f8003, plain, (~ (e1 = op(e4, e1)) | ~ spl130_42), inference(backward_demodulation, [], [f613, f1375])).
fof(f1375, plain, ((e1 = op(e3, e1)) | ~ spl130_42), inference(avatar_component_clause, [], [f1373])).
fof(f8001, plain, (~ spl130_21 | ~ spl130_46), inference(avatar_split_clause, [], [f7992, f1390, f1285])).
fof(f1390, plain, (spl130_46 <=> (e0 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl130_46])])).
fof(f7992, plain, (~ (e0 = op(e4, e0)) | ~ spl130_46), inference(backward_demodulation, [], [f603, f1392])).
fof(f1392, plain, ((e0 = op(e3, e0)) | ~ spl130_46), inference(avatar_component_clause, [], [f1390])).
fof(f7991, plain, (~ spl130_29 | ~ spl130_54), inference(avatar_split_clause, [], [f7988, f1423, f1318])).
fof(f7988, plain, (~ (e3 = op(e3, e4)) | ~ spl130_54), inference(backward_demodulation, [], [f641, f1425])).
fof(f1425, plain, ((e3 = op(e2, e4)) | ~ spl130_54), inference(avatar_component_clause, [], [f1423])).
fof(f641, plain, ~ (op(e2, e4) = op(e3, e4)), inference(cnf_transformation, [], [f4])).
fof(f7987, plain, (~ spl130_53 | ~ spl130_58), inference(avatar_split_clause, [], [f7983, f1440, f1419])).
fof(f7983, plain, (~ (e2 = op(e2, e4)) | ~ spl130_58), inference(backward_demodulation, [], [f673, f1442])).
fof(f1442, plain, ((e2 = op(e2, e3)) | ~ spl130_58), inference(avatar_component_clause, [], [f1440])).
fof(f7980, plain, (~ spl130_12 | ~ spl130_65 | spl130_262), inference(avatar_split_clause, [], [f7973, f2672, f1469, f1247])).
fof(f7973, plain, (~ (e1 = op(e4, e2)) | (~ spl130_65 | spl130_262)), inference(backward_demodulation, [], [f2674, f1471])).
fof(f1471, plain, ((e4 = op(e2, e2)) | ~ spl130_65), inference(avatar_component_clause, [], [f1469])).
fof(f2674, plain, (~ (e1 = op(op(e2, e2), e2)) | spl130_262), inference(avatar_component_clause, [], [f2672])).
fof(f7979, plain, (~ spl130_40 | ~ spl130_65), inference(avatar_split_clause, [], [f7968, f1469, f1364])).
fof(f7968, plain, (~ (e4 = op(e3, e2)) | ~ spl130_65), inference(backward_demodulation, [], [f621, f1471])).
fof(f621, plain, ~ (op(e2, e2) = op(e3, e2)), inference(cnf_transformation, [], [f4])).
fof(f7967, plain, (~ spl130_56 | ~ spl130_66), inference(avatar_split_clause, [], [f7960, f1474, f1432])).
fof(f7960, plain, (~ (e0 = op(e2, e3)) | ~ spl130_66), inference(backward_demodulation, [], [f669, f1476])).
fof(f1476, plain, ((e0 = op(e2, e1)) | ~ spl130_66), inference(avatar_component_clause, [], [f1474])).
fof(f7955, plain, (~ spl130_47 | ~ spl130_72), inference(avatar_split_clause, [], [f7946, f1499, f1394])).
fof(f7946, plain, (~ (e1 = op(e3, e0)) | ~ spl130_72), inference(backward_demodulation, [], [f601, f1501])).
fof(f601, plain, ~ (op(e2, e0) = op(e3, e0)), inference(cnf_transformation, [], [f4])).
fof(f7945, plain, (~ spl130_53 | ~ spl130_78), inference(avatar_split_clause, [], [f7940, f1524, f1419])).
fof(f7940, plain, (~ (e2 = op(e2, e4)) | ~ spl130_78), inference(backward_demodulation, [], [f638, f1526])).
fof(f638, plain, ~ (op(e1, e4) = op(e2, e4)), inference(cnf_transformation, [], [f4])).
fof(f7939, plain, (~ spl130_77 | ~ spl130_82), inference(avatar_split_clause, [], [f7935, f1541, f1520])).
fof(f7935, plain, (~ (e1 = op(e1, e4)) | ~ spl130_82), inference(backward_demodulation, [], [f663, f1543])).
fof(f1543, plain, ((e1 = op(e1, e3)) | ~ spl130_82), inference(avatar_component_clause, [], [f1541])).
fof(f7931, plain, (~ spl130_85 | ~ spl130_100), inference(avatar_split_clause, [], [f7929, f1616, f1553])).
fof(f7929, plain, (~ (e4 = op(e1, e3)) | ~ spl130_100), inference(backward_demodulation, [], [f656, f1618])).
fof(f7917, plain, (~ spl130_101 | ~ spl130_106), inference(avatar_split_clause, [], [f7911, f1642, f1621])).
fof(f1642, plain, (spl130_106 <=> (e0 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl130_106])])).
fof(f7911, plain, (~ (e0 = op(e0, e4)) | ~ spl130_106), inference(backward_demodulation, [], [f653, f1644])).
fof(f1644, plain, ((e0 = op(e0, e3)) | ~ spl130_106), inference(avatar_component_clause, [], [f1642])).
fof(f7916, plain, (~ spl130_56 | ~ spl130_106), inference(avatar_split_clause, [], [f7908, f1642, f1432])).
fof(f7908, plain, (~ (e0 = op(e2, e3)) | ~ spl130_106), inference(backward_demodulation, [], [f625, f1644])).
fof(f7899, plain, (spl130_72 | ~ spl130_123 | ~ spl130_310), inference(avatar_split_clause, [], [f7894, f2918, f1713, f1499])).
fof(f7894, plain, ((e1 = op(e2, e0)) | (~ spl130_123 | ~ spl130_310)), inference(backward_demodulation, [], [f2919, f1715])).
fof(f7897, plain, (~ spl130_118 | ~ spl130_123), inference(avatar_split_clause, [], [f7886, f1713, f1692])).
fof(f7886, plain, (~ (e2 = op(e0, e1)) | ~ spl130_123), inference(backward_demodulation, [], [f644, f1715])).
fof(f7896, plain, (~ spl130_98 | ~ spl130_123), inference(avatar_split_clause, [], [f7882, f1713, f1608])).
fof(f7882, plain, (~ (e2 = op(e1, e0)) | ~ spl130_123), inference(backward_demodulation, [], [f594, f1715])).
fof(f7881, plain, (spl130_10 | ~ spl130_129), inference(avatar_split_clause, [], [f7870, f1738, f1238])).
fof(f1738, plain, (spl130_129 <=> (e3 = unit)), introduced(avatar_definition, [new_symbols(naming, [spl130_129])])).
fof(f7870, plain, ((e4 = op(e4, e3)) | ~ spl130_129), inference(backward_demodulation, [], [f542, f1740])).
fof(f1740, plain, ((e3 = unit) | ~ spl130_129), inference(avatar_component_clause, [], [f1738])).
fof(f7880, plain, (spl130_30 | ~ spl130_129), inference(avatar_split_clause, [], [f7869, f1738, f1322])).
fof(f7869, plain, ((e4 = op(e3, e4)) | ~ spl130_129), inference(backward_demodulation, [], [f541, f1740])).
fof(f7879, plain, (spl130_34 | ~ spl130_129), inference(avatar_split_clause, [], [f7868, f1738, f1339])).
fof(f7868, plain, ((e3 = op(e3, e3)) | ~ spl130_129), inference(backward_demodulation, [], [f540, f1740])).
fof(f7877, plain, (spl130_58 | ~ spl130_129), inference(avatar_split_clause, [], [f7866, f1738, f1440])).
fof(f7866, plain, ((e2 = op(e2, e3)) | ~ spl130_129), inference(backward_demodulation, [], [f538, f1740])).
fof(f7876, plain, (spl130_38 | ~ spl130_129), inference(avatar_split_clause, [], [f7865, f1738, f1356])).
fof(f7865, plain, ((e2 = op(e3, e2)) | ~ spl130_129), inference(backward_demodulation, [], [f537, f1740])).
fof(f7875, plain, (spl130_82 | ~ spl130_129), inference(avatar_split_clause, [], [f7864, f1738, f1541])).
fof(f7864, plain, ((e1 = op(e1, e3)) | ~ spl130_129), inference(backward_demodulation, [], [f536, f1740])).
fof(f7874, plain, (spl130_42 | ~ spl130_129), inference(avatar_split_clause, [], [f7863, f1738, f1373])).
fof(f7863, plain, ((e1 = op(e3, e1)) | ~ spl130_129), inference(backward_demodulation, [], [f535, f1740])).
fof(f7873, plain, (spl130_106 | ~ spl130_129), inference(avatar_split_clause, [], [f7862, f1738, f1642])).
fof(f7862, plain, ((e0 = op(e0, e3)) | ~ spl130_129), inference(backward_demodulation, [], [f534, f1740])).
fof(f7872, plain, (spl130_46 | ~ spl130_129), inference(avatar_split_clause, [], [f7861, f1738, f1390])).
fof(f7861, plain, ((e0 = op(e3, e0)) | ~ spl130_129), inference(backward_demodulation, [], [f533, f1740])).
fof(f7840, plain, (~ spl130_45 | ~ spl130_94 | spl130_286), inference(avatar_split_clause, [], [f7776, f2786, f1591, f1385])).
fof(f7776, plain, (~ (e4 = op(e3, e1)) | (~ spl130_94 | spl130_286)), inference(backward_demodulation, [], [f2788, f1593])).
fof(f2788, plain, (~ (e4 = op(op(e1, e1), e1)) | spl130_286), inference(avatar_component_clause, [], [f2786])).
fof(f7808, plain, (~ spl130_69 | ~ spl130_74), inference(avatar_split_clause, [], [f7804, f1507, f1486])).
fof(f1486, plain, (spl130_69 <=> (e3 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl130_69])])).
fof(f7804, plain, (~ (e3 = op(e2, e1)) | ~ spl130_74), inference(backward_demodulation, [], [f664, f1509])).
fof(f7799, plain, (~ spl130_81 | ~ spl130_86), inference(avatar_split_clause, [], [f7793, f1558, f1537])).
fof(f7793, plain, (~ (e0 = op(e1, e3)) | ~ spl130_86), inference(backward_demodulation, [], [f661, f1560])).
fof(f7785, plain, (~ spl130_33 | ~ spl130_94 | spl130_282), inference(avatar_split_clause, [], [f7775, f2765, f1591, f1335])).
fof(f7775, plain, (~ (e2 = op(e3, e3)) | (~ spl130_94 | spl130_282)), inference(backward_demodulation, [], [f2767, f1593])).
fof(f7784, plain, (~ spl130_35 | ~ spl130_94 | spl130_275), inference(avatar_split_clause, [], [f7774, f2730, f1591, f1343])).
fof(f2730, plain, (spl130_275 <=> (e4 = op(op(e1, e1), op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl130_275])])).
fof(f7774, plain, (~ (e4 = op(e3, e3)) | (~ spl130_94 | spl130_275)), inference(backward_demodulation, [], [f2732, f1593])).
fof(f2732, plain, (~ (e4 = op(op(e1, e1), op(e1, e1))) | spl130_275), inference(avatar_component_clause, [], [f2730])).
fof(f7781, plain, (~ spl130_69 | ~ spl130_94), inference(avatar_split_clause, [], [f7768, f1591, f1486])).
fof(f7768, plain, (~ (e3 = op(e2, e1)) | ~ spl130_94), inference(backward_demodulation, [], [f608, f1593])).
fof(f7767, plain, (~ spl130_93 | ~ spl130_98), inference(avatar_split_clause, [], [f7761, f1608, f1587])).
fof(f7761, plain, (~ (e2 = op(e1, e1)) | ~ spl130_98), inference(backward_demodulation, [], [f654, f1610])).
fof(f7766, plain, (~ spl130_48 | ~ spl130_98), inference(avatar_split_clause, [], [f7760, f1608, f1398])).
fof(f7760, plain, (~ (e2 = op(e3, e0)) | ~ spl130_98), inference(backward_demodulation, [], [f599, f1610])).
fof(f7758, plain, (~ spl130_57 | ~ spl130_107), inference(avatar_split_clause, [], [f7753, f1646, f1436])).
fof(f7753, plain, (~ (e1 = op(e2, e3)) | ~ spl130_107), inference(backward_demodulation, [], [f625, f1648])).
fof(f7750, plain, (~ spl130_93 | ~ spl130_118), inference(avatar_split_clause, [], [f7745, f1692, f1587])).
fof(f7745, plain, (~ (e2 = op(e1, e1)) | ~ spl130_118), inference(backward_demodulation, [], [f604, f1694])).
fof(f604, plain, ~ (op(e0, e1) = op(e1, e1)), inference(cnf_transformation, [], [f4])).
fof(f7737, plain, (~ spl130_75 | ~ spl130_125), inference(avatar_split_clause, [], [f7729, f1721, f1511])).
fof(f7729, plain, (~ (e4 = op(e2, e0)) | ~ spl130_125), inference(backward_demodulation, [], [f595, f1723])).
fof(f595, plain, ~ (op(e0, e0) = op(e2, e0)), inference(cnf_transformation, [], [f4])).
fof(f7723, plain, (~ spl130_125 | ~ spl130_290 | spl130_318), inference(avatar_split_clause, [], [f7717, f2994, f2808, f1721])).
fof(f2808, plain, (spl130_290 <=> (e0 = op(op(e2, e2), e2))), introduced(avatar_definition, [new_symbols(naming, [spl130_290])])).
fof(f2994, plain, (spl130_318 <=> (e4 = op(op(op(e2, e2), e2), op(op(e2, e2), e2)))), introduced(avatar_definition, [new_symbols(naming, [spl130_318])])).
fof(f7717, plain, (~ (op(e0, e0) = e4) | (~ spl130_290 | spl130_318)), inference(backward_demodulation, [], [f2996, f2809])).
fof(f2809, plain, ((e0 = op(op(e2, e2), e2)) | ~ spl130_290), inference(avatar_component_clause, [], [f2808])).
fof(f2996, plain, (~ (e4 = op(op(op(e2, e2), e2), op(op(e2, e2), e2))) | spl130_318), inference(avatar_component_clause, [], [f2994])).
fof(f7716, plain, (~ spl130_125 | ~ spl130_292 | spl130_319), inference(avatar_split_clause, [], [f7711, f3000, f2818, f1721])).
fof(f2818, plain, (spl130_292 <=> (e0 = op(op(e3, e3), e3))), introduced(avatar_definition, [new_symbols(naming, [spl130_292])])).
fof(f7711, plain, (~ (op(e0, e0) = e4) | (~ spl130_292 | spl130_319)), inference(backward_demodulation, [], [f3002, f2819])).
fof(f2819, plain, ((e0 = op(op(e3, e3), e3)) | ~ spl130_292), inference(avatar_component_clause, [], [f2818])).
fof(f7700, plain, (~ spl130_33 | ~ spl130_301 | spl130_309), inference(avatar_split_clause, [], [f7698, f2914, f2864, f1335])).
fof(f7698, plain, (~ (e2 = op(e3, e3)) | (~ spl130_301 | spl130_309)), inference(backward_demodulation, [], [f2916, f2865])).
fof(f7662, plain, (~ spl130_3 | ~ spl130_35 | spl130_268), inference(avatar_split_clause, [], [f7661, f2699, f1343, f1209])).
fof(f2699, plain, (spl130_268 <=> (e2 = op(op(e3, e3), op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl130_268])])).
fof(f7661, plain, (~ (e2 = op(e4, e4)) | (~ spl130_35 | spl130_268)), inference(forward_demodulation, [], [f2701, f1345])).
fof(f1345, plain, ((e4 = op(e3, e3)) | ~ spl130_35), inference(avatar_component_clause, [], [f1343])).
fof(f2701, plain, (~ (e2 = op(op(e3, e3), op(e3, e3))) | spl130_268), inference(avatar_component_clause, [], [f2699])).
fof(f7646, plain, (~ spl130_35 | ~ spl130_99 | ~ spl130_122 | spl130_320), inference(avatar_contradiction_clause, [], [f7645])).
fof(f7645, plain, ($false | (~ spl130_35 | ~ spl130_99 | ~ spl130_122 | spl130_320)), inference(subsumption_resolution, [], [f7644, f1345])).
fof(f7644, plain, (~ (e4 = op(e3, e3)) | (~ spl130_99 | ~ spl130_122 | spl130_320)), inference(forward_demodulation, [], [f7643, f1614])).
fof(f7643, plain, (~ (e4 = op(op(e1, e0), op(e1, e0))) | (~ spl130_122 | spl130_320)), inference(forward_demodulation, [], [f3008, f1711])).
fof(f7595, plain, (~ spl130_6 | ~ spl130_21), inference(avatar_split_clause, [], [f7590, f1285, f1222])).
fof(f7590, plain, (~ (e0 = op(e4, e3)) | ~ spl130_21), inference(backward_demodulation, [], [f686, f1287])).
fof(f1287, plain, ((e0 = op(e4, e0)) | ~ spl130_21), inference(avatar_component_clause, [], [f1285])).
fof(f7565, plain, (~ spl130_51 | ~ spl130_61), inference(avatar_split_clause, [], [f7556, f1453, f1411])).
fof(f1411, plain, (spl130_51 <=> (e0 = op(e2, e4))), introduced(avatar_definition, [new_symbols(naming, [spl130_51])])).
fof(f7556, plain, (~ (e0 = op(e2, e4)) | ~ spl130_61), inference(backward_demodulation, [], [f672, f1455])).
fof(f7551, plain, (~ spl130_64 | ~ spl130_69), inference(avatar_split_clause, [], [f7544, f1486, f1465])).
fof(f7544, plain, (~ (e3 = op(e2, e2)) | ~ spl130_69), inference(backward_demodulation, [], [f668, f1488])).
fof(f1488, plain, ((e3 = op(e2, e1)) | ~ spl130_69), inference(avatar_component_clause, [], [f1486])).
fof(f7550, plain, (~ spl130_44 | ~ spl130_69), inference(avatar_split_clause, [], [f7542, f1486, f1381])).
fof(f7542, plain, (~ (e3 = op(e3, e1)) | ~ spl130_69), inference(backward_demodulation, [], [f611, f1488])).
fof(f7541, plain, (~ spl130_27 | ~ spl130_77), inference(avatar_split_clause, [], [f7537, f1520, f1310])).
fof(f7537, plain, (~ (e1 = op(e3, e4)) | ~ spl130_77), inference(backward_demodulation, [], [f639, f1522])).
fof(f1522, plain, ((e1 = op(e1, e4)) | ~ spl130_77), inference(avatar_component_clause, [], [f1520])).
fof(f7522, plain, (spl130_61 | ~ spl130_93 | ~ spl130_312), inference(avatar_split_clause, [], [f7516, f2935, f1587, f1453])).
fof(f7516, plain, ((e0 = op(e2, e2)) | (~ spl130_93 | ~ spl130_312)), inference(backward_demodulation, [], [f2936, f1589])).
fof(f2936, plain, ((e0 = op(op(e1, e1), op(e1, e1))) | ~ spl130_312), inference(avatar_component_clause, [], [f2935])).
fof(f7501, plain, (~ spl130_84 | ~ spl130_99), inference(avatar_split_clause, [], [f7496, f1612, f1549])).
fof(f7496, plain, (~ (e3 = op(e1, e3)) | ~ spl130_99), inference(backward_demodulation, [], [f656, f1614])).
fof(f7484, plain, (~ spl130_104 | ~ spl130_114), inference(avatar_split_clause, [], [f7481, f1675, f1633])).
fof(f7481, plain, (~ (e3 = op(e0, e4)) | ~ spl130_114), inference(backward_demodulation, [], [f652, f1677])).
fof(f1677, plain, ((e3 = op(e0, e2)) | ~ spl130_114), inference(avatar_component_clause, [], [f1675])).
fof(f7483, plain, (~ spl130_64 | ~ spl130_114), inference(avatar_split_clause, [], [f7478, f1675, f1465])).
fof(f7478, plain, (~ (e3 = op(e2, e2)) | ~ spl130_114), inference(backward_demodulation, [], [f615, f1677])).
fof(f7475, plain, (~ spl130_20 | ~ spl130_120), inference(avatar_split_clause, [], [f7472, f1700, f1280])).
fof(f7472, plain, (~ (e4 = op(e4, e1)) | ~ spl130_120), inference(backward_demodulation, [], [f607, f1702])).
fof(f7441, plain, (~ spl130_64 | ~ spl130_3 | spl130_271), inference(avatar_split_clause, [], [f7440, f2712, f1209, f1465])).
fof(f7440, plain, (~ (e3 = op(e2, e2)) | (~ spl130_3 | spl130_271)), inference(forward_demodulation, [], [f2714, f1211])).
fof(f7439, plain, (~ spl130_35 | ~ spl130_283 | spl130_318), inference(avatar_contradiction_clause, [], [f7438])).
fof(f7438, plain, ($false | (~ spl130_35 | ~ spl130_283 | spl130_318)), inference(subsumption_resolution, [], [f7437, f1345])).
fof(f7437, plain, (~ (e4 = op(e3, e3)) | (~ spl130_283 | spl130_318)), inference(backward_demodulation, [], [f2996, f2771])).
fof(f7413, plain, (~ spl130_100 | ~ spl130_75), inference(avatar_split_clause, [], [f7412, f1511, f1616])).
fof(f7412, plain, (~ (e4 = op(e1, e0)) | ~ spl130_75), inference(forward_demodulation, [], [f598, f1513])).
fof(f7400, plain, (~ spl130_61 | ~ spl130_3 | spl130_300), inference(avatar_split_clause, [], [f7361, f2859, f1209, f1453])).
fof(f2859, plain, (spl130_300 <=> (e0 = op(op(e4, e4), op(e4, e4)))), introduced(avatar_definition, [new_symbols(naming, [spl130_300])])).
fof(f7361, plain, (~ (e0 = op(e2, e2)) | (~ spl130_3 | spl130_300)), inference(forward_demodulation, [], [f2861, f1211])).
fof(f2861, plain, (~ (e0 = op(op(e4, e4), op(e4, e4))) | spl130_300), inference(avatar_component_clause, [], [f2859])).
fof(f7395, plain, (~ spl130_55 | ~ spl130_75), inference(avatar_split_clause, [], [f7394, f1511, f1427])).
fof(f7394, plain, (~ (e4 = op(e2, e4)) | ~ spl130_75), inference(forward_demodulation, [], [f667, f1513])).
fof(f7393, plain, (~ spl130_52 | ~ spl130_3 | spl130_270), inference(avatar_split_clause, [], [f7335, f2708, f1209, f1415])).
fof(f7335, plain, (~ (e1 = op(e2, e4)) | (~ spl130_3 | spl130_270)), inference(backward_demodulation, [], [f2710, f1211])).
fof(f2710, plain, (~ (e1 = op(op(e4, e4), e4)) | spl130_270), inference(avatar_component_clause, [], [f2708])).
fof(f7366, plain, (~ spl130_3 | ~ spl130_51 | ~ spl130_122 | spl130_293), inference(avatar_contradiction_clause, [], [f7365])).
fof(f7365, plain, ($false | (~ spl130_3 | ~ spl130_51 | ~ spl130_122 | spl130_293)), inference(subsumption_resolution, [], [f7364, f1711])).
fof(f7364, plain, (~ (op(e0, e0) = e1) | (~ spl130_3 | ~ spl130_51 | spl130_293)), inference(forward_demodulation, [], [f7363, f1413])).
fof(f1413, plain, ((e0 = op(e2, e4)) | ~ spl130_51), inference(avatar_component_clause, [], [f1411])).
fof(f7363, plain, (~ (e1 = op(op(e2, e4), op(e2, e4))) | (~ spl130_3 | spl130_293)), inference(forward_demodulation, [], [f2826, f1211])).
fof(f7342, plain, (~ spl130_3 | ~ spl130_51 | spl130_294), inference(avatar_contradiction_clause, [], [f7341])).
fof(f7341, plain, ($false | (~ spl130_3 | ~ spl130_51 | spl130_294)), inference(subsumption_resolution, [], [f7337, f1413])).
fof(f7337, plain, (~ (e0 = op(e2, e4)) | (~ spl130_3 | spl130_294)), inference(backward_demodulation, [], [f2830, f1211])).
fof(f2830, plain, (~ (e0 = op(op(e4, e4), e4)) | spl130_294), inference(avatar_component_clause, [], [f2828])).
fof(f7321, plain, (~ spl130_4 | ~ spl130_24), inference(avatar_split_clause, [], [f7319, f1297, f1213])).
fof(f7319, plain, (~ (e3 = op(e4, e4)) | ~ spl130_24), inference(backward_demodulation, [], [f687, f1299])).
fof(f687, plain, ~ (op(e4, e0) = op(e4, e4)), inference(cnf_transformation, [], [f4])).
fof(f7312, plain, (~ spl130_32 | ~ spl130_35), inference(avatar_contradiction_clause, [], [f7311])).
fof(f7311, plain, ($false | (~ spl130_32 | ~ spl130_35)), inference(subsumption_resolution, [], [f7310, f700])).
fof(f700, plain, ~ (e1 = e4), inference(cnf_transformation, [], [f5])).
fof(f7310, plain, ((e1 = e4) | (~ spl130_32 | ~ spl130_35)), inference(backward_demodulation, [], [f1345, f1333])).
fof(f7302, plain, (~ spl130_7 | ~ spl130_35 | spl130_266), inference(avatar_split_clause, [], [f7297, f2690, f1343, f1226])).
fof(f7297, plain, (~ (e1 = op(e4, e3)) | (~ spl130_35 | spl130_266)), inference(backward_demodulation, [], [f2692, f1345])).
fof(f2692, plain, (~ (e1 = op(op(e3, e3), e3)) | spl130_266), inference(avatar_component_clause, [], [f2690])).
fof(f7293, plain, (~ spl130_31 | ~ spl130_36), inference(avatar_split_clause, [], [f7284, f1348, f1327])).
fof(f7284, plain, (~ (e0 = op(e3, e3)) | ~ spl130_36), inference(backward_demodulation, [], [f681, f1350])).
fof(f7267, plain, (~ spl130_36 | ~ spl130_64 | spl130_290), inference(avatar_split_clause, [], [f7264, f2808, f1465, f1348])).
fof(f7264, plain, (~ (e0 = op(e3, e2)) | (~ spl130_64 | spl130_290)), inference(backward_demodulation, [], [f2810, f1467])).
fof(f2810, plain, (~ (e0 = op(op(e2, e2), e2)) | spl130_290), inference(avatar_component_clause, [], [f2808])).
fof(f7255, plain, (~ spl130_60 | ~ spl130_75), inference(avatar_split_clause, [], [f7254, f1511, f1448])).
fof(f7254, plain, (~ (e4 = op(e2, e3)) | ~ spl130_75), inference(backward_demodulation, [], [f666, f1513])).
fof(f666, plain, ~ (op(e2, e0) = op(e2, e3)), inference(cnf_transformation, [], [f4])).
fof(f7235, plain, (~ spl130_47 | ~ spl130_122), inference(avatar_split_clause, [], [f7223, f1709, f1394])).
fof(f7223, plain, (~ (e1 = op(e3, e0)) | ~ spl130_122), inference(backward_demodulation, [], [f596, f1711])).
fof(f7192, plain, (~ spl130_34 | ~ spl130_44), inference(avatar_split_clause, [], [f7191, f1381, f1339])).
fof(f7191, plain, (~ (e3 = op(e3, e3)) | ~ spl130_44), inference(forward_demodulation, [], [f679, f1383])).
fof(f1383, plain, ((e3 = op(e3, e1)) | ~ spl130_44), inference(avatar_component_clause, [], [f1381])).
fof(f7190, plain, (~ spl130_34 | ~ spl130_84), inference(avatar_split_clause, [], [f7189, f1549, f1339])).
fof(f7189, plain, (~ (e3 = op(e3, e3)) | ~ spl130_84), inference(forward_demodulation, [], [f629, f1551])).
fof(f1551, plain, ((e3 = op(e1, e3)) | ~ spl130_84), inference(avatar_component_clause, [], [f1549])).
fof(f7188, plain, (~ spl130_33 | ~ spl130_108), inference(avatar_split_clause, [], [f7187, f1650, f1335])).
fof(f7187, plain, (~ (e2 = op(e3, e3)) | ~ spl130_108), inference(forward_demodulation, [], [f626, f1652])).
fof(f7146, plain, (~ spl130_29 | ~ spl130_44), inference(avatar_split_clause, [], [f7144, f1381, f1318])).
fof(f7144, plain, (~ (e3 = op(e3, e4)) | ~ spl130_44), inference(backward_demodulation, [], [f680, f1383])).
fof(f7115, plain, (~ spl130_53 | ~ spl130_68), inference(avatar_split_clause, [], [f7111, f1482, f1419])).
fof(f7111, plain, (~ (e2 = op(e2, e4)) | ~ spl130_68), inference(backward_demodulation, [], [f670, f1484])).
fof(f1484, plain, ((e2 = op(e2, e1)) | ~ spl130_68), inference(avatar_component_clause, [], [f1482])).
fof(f670, plain, ~ (op(e2, e1) = op(e2, e4)), inference(cnf_transformation, [], [f4])).
fof(f7114, plain, (~ spl130_43 | ~ spl130_68), inference(avatar_split_clause, [], [f7108, f1482, f1377])).
fof(f7108, plain, (~ (e2 = op(e3, e1)) | ~ spl130_68), inference(backward_demodulation, [], [f611, f1484])).
fof(f7103, plain, (~ spl130_9 | ~ spl130_84), inference(avatar_split_clause, [], [f7100, f1549, f1234])).
fof(f7100, plain, (~ (e3 = op(e4, e3)) | ~ spl130_84), inference(backward_demodulation, [], [f630, f1551])).
fof(f7099, plain, (~ spl130_83 | ~ spl130_88), inference(avatar_split_clause, [], [f7094, f1566, f1545])).
fof(f7094, plain, (~ (e2 = op(e1, e3)) | ~ spl130_88), inference(backward_demodulation, [], [f661, f1568])).
fof(f1568, plain, ((e2 = op(e1, e2)) | ~ spl130_88), inference(avatar_component_clause, [], [f1566])).
fof(f7098, plain, (~ spl130_13 | ~ spl130_88), inference(avatar_split_clause, [], [f7093, f1566, f1251])).
fof(f7093, plain, (~ (e2 = op(e4, e2)) | ~ spl130_88), inference(backward_demodulation, [], [f620, f1568])).
fof(f7089, plain, (~ spl130_21 | ~ spl130_96), inference(avatar_split_clause, [], [f7081, f1600, f1285])).
fof(f7081, plain, (~ (e0 = op(e4, e0)) | ~ spl130_96), inference(backward_demodulation, [], [f600, f1602])).
fof(f7073, plain, (~ spl130_89 | ~ spl130_114), inference(avatar_split_clause, [], [f7067, f1675, f1570])).
fof(f7067, plain, (~ (e3 = op(e1, e2)) | ~ spl130_114), inference(backward_demodulation, [], [f614, f1677])).
fof(f7066, plain, (~ spl130_111 | ~ spl130_116), inference(avatar_split_clause, [], [f7060, f1684, f1663])).
fof(f7060, plain, (~ (e0 = op(e0, e2)) | ~ spl130_116), inference(backward_demodulation, [], [f648, f1686])).
fof(f7055, plain, (~ spl130_23 | ~ spl130_125 | spl130_296), inference(avatar_split_clause, [], [f7046, f2838, f1721, f1293])).
fof(f7046, plain, (~ (e2 = op(e4, e0)) | (~ spl130_125 | spl130_296)), inference(backward_demodulation, [], [f2840, f1723])).
fof(f2840, plain, (~ (e2 = op(op(e0, e0), e0)) | spl130_296), inference(avatar_component_clause, [], [f2838])).
fof(f6975, plain, (~ spl130_77 | ~ spl130_92), inference(avatar_split_clause, [], [f6974, f1583, f1520])).
fof(f6974, plain, (~ (e1 = op(e1, e4)) | ~ spl130_92), inference(forward_demodulation, [], [f660, f1585])).
fof(f1585, plain, ((e1 = op(e1, e1)) | ~ spl130_92), inference(avatar_component_clause, [], [f1583])).
fof(f6962, plain, (~ spl130_124 | ~ spl130_74), inference(avatar_split_clause, [], [f6961, f1507, f1717])).
fof(f6961, plain, (~ (op(e0, e0) = e3) | ~ spl130_74), inference(forward_demodulation, [], [f595, f1509])).
fof(f6956, plain, (~ spl130_65 | ~ spl130_60), inference(avatar_split_clause, [], [f6955, f1448, f1469])).
fof(f6955, plain, (~ (e4 = op(e2, e2)) | ~ spl130_60), inference(forward_demodulation, [], [f671, f1450])).
fof(f6947, plain, (~ spl130_110 | ~ spl130_31 | spl130_288), inference(avatar_split_clause, [], [f6946, f2798, f1327, f1658])).
fof(f6946, plain, (~ (e4 = op(e0, e3)) | (~ spl130_31 | spl130_288)), inference(forward_demodulation, [], [f2800, f1329])).
fof(f1329, plain, ((e0 = op(e3, e3)) | ~ spl130_31), inference(avatar_component_clause, [], [f1327])).
fof(f2800, plain, (~ (e4 = op(op(e3, e3), e3)) | spl130_288), inference(avatar_component_clause, [], [f2798])).
fof(f6916, plain, (~ spl130_11 | ~ spl130_21), inference(avatar_split_clause, [], [f6909, f1285, f1243])).
fof(f6909, plain, (~ (e0 = op(e4, e2)) | ~ spl130_21), inference(backward_demodulation, [], [f685, f1287])).
fof(f6905, plain, (~ spl130_31 | ~ spl130_35), inference(avatar_contradiction_clause, [], [f6904])).
fof(f6904, plain, ($false | (~ spl130_31 | ~ spl130_35)), inference(subsumption_resolution, [], [f6903, f697])).
fof(f697, plain, ~ (e0 = e4), inference(cnf_transformation, [], [f5])).
fof(f6903, plain, ((e0 = e4) | (~ spl130_31 | ~ spl130_35)), inference(backward_demodulation, [], [f1345, f1329])).
fof(f6870, plain, (~ spl130_42 | ~ spl130_47), inference(avatar_split_clause, [], [f6863, f1394, f1373])).
fof(f6863, plain, (~ (e1 = op(e3, e1)) | ~ spl130_47), inference(backward_demodulation, [], [f674, f1396])).
fof(f6856, plain, (~ spl130_35 | ~ spl130_60), inference(avatar_split_clause, [], [f6853, f1448, f1343])).
fof(f6853, plain, (~ (e4 = op(e3, e3)) | ~ spl130_60), inference(backward_demodulation, [], [f631, f1450])).
fof(f6852, plain, (~ spl130_72 | ~ spl130_74), inference(avatar_contradiction_clause, [], [f6851])).
fof(f6851, plain, ($false | (~ spl130_72 | ~ spl130_74)), inference(subsumption_resolution, [], [f6850, f699])).
fof(f699, plain, ~ (e1 = e3), inference(cnf_transformation, [], [f5])).
fof(f6850, plain, ((e1 = e3) | (~ spl130_72 | ~ spl130_74)), inference(forward_demodulation, [], [f1509, f1501])).
fof(f6800, plain, (~ spl130_106 | ~ spl130_111), inference(avatar_split_clause, [], [f6793, f1663, f1642])).
fof(f6793, plain, (~ (e0 = op(e0, e3)) | ~ spl130_111), inference(backward_demodulation, [], [f651, f1665])).
fof(f6786, plain, (~ spl130_72 | ~ spl130_123 | spl130_310), inference(avatar_contradiction_clause, [], [f6785])).
fof(f6785, plain, ($false | (~ spl130_72 | ~ spl130_123 | spl130_310)), inference(subsumption_resolution, [], [f6774, f1501])).
fof(f6774, plain, (~ (e1 = op(e2, e0)) | (~ spl130_123 | spl130_310)), inference(backward_demodulation, [], [f2920, f1715])).
fof(f6777, plain, (~ spl130_23 | ~ spl130_123), inference(avatar_split_clause, [], [f6763, f1713, f1293])).
fof(f6763, plain, (~ (e2 = op(e4, e0)) | ~ spl130_123), inference(backward_demodulation, [], [f597, f1715])).
fof(f6761, plain, (spl130_5 | ~ spl130_130), inference(avatar_split_clause, [], [f6748, f1742, f1217])).
fof(f6748, plain, ((e4 = op(e4, e4)) | ~ spl130_130), inference(backward_demodulation, [], [f542, f1744])).
fof(f6755, plain, (~ spl130_76 | ~ spl130_130), inference(avatar_contradiction_clause, [], [f6754])).
fof(f6754, plain, ($false | (~ spl130_76 | ~ spl130_130)), inference(subsumption_resolution, [], [f6753, f694])).
fof(f694, plain, ~ (e0 = e1), inference(cnf_transformation, [], [f5])).
fof(f6753, plain, ((e0 = e1) | (~ spl130_76 | ~ spl130_130)), inference(forward_demodulation, [], [f6742, f1518])).
fof(f1518, plain, ((e0 = op(e1, e4)) | ~ spl130_76), inference(avatar_component_clause, [], [f1516])).
fof(f6742, plain, ((e1 = op(e1, e4)) | ~ spl130_130), inference(backward_demodulation, [], [f536, f1744])).
fof(f6737, plain, (~ spl130_11 | ~ spl130_65 | spl130_290), inference(avatar_split_clause, [], [f6736, f2808, f1469, f1243])).
fof(f6736, plain, (~ (e0 = op(e4, e2)) | (~ spl130_65 | spl130_290)), inference(forward_demodulation, [], [f2810, f1471])).
fof(f6728, plain, (~ spl130_110 | ~ spl130_120), inference(avatar_split_clause, [], [f6727, f1700, f1658])).
fof(f6727, plain, (~ (e4 = op(e0, e3)) | ~ spl130_120), inference(forward_demodulation, [], [f649, f1702])).
fof(f649, plain, ~ (op(e0, e1) = op(e0, e3)), inference(cnf_transformation, [], [f4])).
fof(f6721, plain, (~ spl130_84 | ~ spl130_89), inference(avatar_split_clause, [], [f6720, f1570, f1549])).
fof(f6720, plain, (~ (e3 = op(e1, e3)) | ~ spl130_89), inference(forward_demodulation, [], [f661, f1572])).
fof(f6701, plain, (~ spl130_14 | ~ spl130_89), inference(avatar_split_clause, [], [f6700, f1570, f1255])).
fof(f6700, plain, (~ (e3 = op(e4, e2)) | ~ spl130_89), inference(forward_demodulation, [], [f620, f1572])).
fof(f6697, plain, (~ spl130_2 | ~ spl130_76 | spl130_294), inference(avatar_contradiction_clause, [], [f6696])).
fof(f6696, plain, ($false | (~ spl130_2 | ~ spl130_76 | spl130_294)), inference(subsumption_resolution, [], [f6689, f1518])).
fof(f6689, plain, (~ (e0 = op(e1, e4)) | (~ spl130_2 | spl130_294)), inference(backward_demodulation, [], [f2830, f1207])).
fof(f6672, plain, (~ spl130_4 | ~ spl130_54), inference(avatar_split_clause, [], [f6670, f1423, f1213])).
fof(f6670, plain, (~ (e3 = op(e4, e4)) | ~ spl130_54), inference(backward_demodulation, [], [f642, f1425])).
fof(f642, plain, ~ (op(e2, e4) = op(e4, e4)), inference(cnf_transformation, [], [f4])).
fof(f6669, plain, (~ spl130_2 | ~ spl130_65 | spl130_284), inference(avatar_split_clause, [], [f6667, f2775, f1469, f1205])).
fof(f6667, plain, (~ (e1 = op(e4, e4)) | (~ spl130_65 | spl130_284)), inference(backward_demodulation, [], [f2777, f1471])).
fof(f6664, plain, (~ spl130_51 | ~ spl130_66), inference(avatar_split_clause, [], [f6659, f1474, f1411])).
fof(f6659, plain, (~ (e0 = op(e2, e4)) | ~ spl130_66), inference(backward_demodulation, [], [f670, f1476])).
fof(f6655, plain, (~ spl130_22 | ~ spl130_72), inference(avatar_split_clause, [], [f6648, f1499, f1289])).
fof(f6648, plain, (~ (e1 = op(e4, e0)) | ~ spl130_72), inference(backward_demodulation, [], [f602, f1501])).
fof(f6637, plain, (spl130_66 | ~ spl130_93 | ~ spl130_306), inference(avatar_split_clause, [], [f6631, f2896, f1587, f1474])).
fof(f6631, plain, ((e0 = op(e2, e1)) | (~ spl130_93 | ~ spl130_306)), inference(backward_demodulation, [], [f2897, f1589])).
fof(f2897, plain, ((e0 = op(op(e1, e1), e1)) | ~ spl130_306), inference(avatar_component_clause, [], [f2896])).
fof(f6634, plain, (~ spl130_78 | ~ spl130_93), inference(avatar_split_clause, [], [f6627, f1587, f1524])).
fof(f6627, plain, (~ (e2 = op(e1, e4)) | ~ spl130_93), inference(backward_demodulation, [], [f660, f1589])).
fof(f6633, plain, (~ spl130_18 | ~ spl130_93), inference(avatar_split_clause, [], [f6625, f1587, f1272])).
fof(f6625, plain, (~ (e2 = op(e4, e1)) | ~ spl130_93), inference(backward_demodulation, [], [f610, f1589])).
fof(f610, plain, ~ (op(e1, e1) = op(e4, e1)), inference(cnf_transformation, [], [f4])).
fof(f6623, plain, (~ spl130_90 | ~ spl130_100), inference(avatar_split_clause, [], [f6620, f1616, f1574])).
fof(f6620, plain, (~ (e4 = op(e1, e2)) | ~ spl130_100), inference(backward_demodulation, [], [f655, f1618])).
fof(f6610, plain, (~ spl130_102 | ~ spl130_112), inference(avatar_split_clause, [], [f6605, f1667, f1625])).
fof(f6605, plain, (~ (e1 = op(e0, e4)) | ~ spl130_112), inference(backward_demodulation, [], [f652, f1669])).
fof(f6591, plain, (~ spl130_114 | ~ spl130_124), inference(avatar_split_clause, [], [f6577, f1717, f1675])).
fof(f6577, plain, (~ (e3 = op(e0, e2)) | ~ spl130_124), inference(backward_demodulation, [], [f645, f1719])).
fof(f6553, plain, (~ spl130_55 | ~ spl130_30), inference(avatar_split_clause, [], [f6411, f1322, f1427])).
fof(f6411, plain, (~ (e4 = op(e2, e4)) | ~ spl130_30), inference(forward_demodulation, [], [f641, f1324])).
fof(f1324, plain, ((e4 = op(e3, e4)) | ~ spl130_30), inference(avatar_component_clause, [], [f1322])).
fof(f6529, plain, (~ spl130_4 | ~ spl130_75 | ~ spl130_123 | spl130_316), inference(avatar_contradiction_clause, [], [f6528])).
fof(f6528, plain, ($false | (~ spl130_4 | ~ spl130_75 | ~ spl130_123 | spl130_316)), inference(subsumption_resolution, [], [f6527, f1215])).
fof(f6527, plain, (~ (e3 = op(e4, e4)) | (~ spl130_75 | ~ spl130_123 | spl130_316)), inference(forward_demodulation, [], [f6526, f1513])).
fof(f6526, plain, (~ (e3 = op(op(e2, e0), op(e2, e0))) | (~ spl130_123 | spl130_316)), inference(forward_demodulation, [], [f2968, f1715])).
fof(f6506, plain, (~ spl130_62 | ~ spl130_91 | spl130_303), inference(avatar_contradiction_clause, [], [f6505])).
fof(f6505, plain, ($false | (~ spl130_62 | ~ spl130_91 | spl130_303)), inference(subsumption_resolution, [], [f6501, f1581])).
fof(f6501, plain, (~ (e0 = op(e1, e1)) | (~ spl130_62 | spl130_303)), inference(backward_demodulation, [], [f2877, f1459])).
fof(f2877, plain, (~ (e0 = op(op(e2, e2), op(e2, e2))) | spl130_303), inference(avatar_component_clause, [], [f2875])).
fof(f2875, plain, (spl130_303 <=> (e0 = op(op(e2, e2), op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl130_303])])).
fof(f6441, plain, (~ spl130_75 | ~ spl130_123 | spl130_304), inference(avatar_contradiction_clause, [], [f6440])).
fof(f6440, plain, ($false | (~ spl130_75 | ~ spl130_123 | spl130_304)), inference(subsumption_resolution, [], [f6431, f1513])).
fof(f6431, plain, (~ (e4 = op(e2, e0)) | (~ spl130_123 | spl130_304)), inference(backward_demodulation, [], [f2884, f1715])).
fof(f2884, plain, (~ (e4 = op(op(e0, e0), e0)) | spl130_304), inference(avatar_component_clause, [], [f2882])).
fof(f6391, plain, (~ spl130_82 | ~ spl130_85), inference(avatar_contradiction_clause, [], [f6390])).
fof(f6390, plain, ($false | (~ spl130_82 | ~ spl130_85)), inference(subsumption_resolution, [], [f6389, f700])).
fof(f6389, plain, ((e1 = e4) | (~ spl130_82 | ~ spl130_85)), inference(forward_demodulation, [], [f1555, f1543])).
fof(f1555, plain, ((e4 = op(e1, e3)) | ~ spl130_85), inference(avatar_component_clause, [], [f1553])).
fof(f6382, plain, (~ spl130_61 | ~ spl130_93 | spl130_312), inference(avatar_split_clause, [], [f6379, f2935, f1587, f1453])).
fof(f6379, plain, (~ (e0 = op(e2, e2)) | (~ spl130_93 | spl130_312)), inference(backward_demodulation, [], [f2937, f1589])).
fof(f6347, plain, (~ spl130_100 | ~ spl130_122 | spl130_304), inference(avatar_split_clause, [], [f6337, f2882, f1709, f1616])).
fof(f6337, plain, (~ (e4 = op(e1, e0)) | (~ spl130_122 | spl130_304)), inference(backward_demodulation, [], [f2884, f1711])).
fof(f6345, plain, (~ spl130_99 | ~ spl130_122 | spl130_301), inference(avatar_split_clause, [], [f6335, f2864, f1709, f1612])).
fof(f6335, plain, (~ (e3 = op(e1, e0)) | (~ spl130_122 | spl130_301)), inference(backward_demodulation, [], [f2866, f1711])).
fof(f2866, plain, (~ (e3 = op(op(e0, e0), e0)) | spl130_301), inference(avatar_component_clause, [], [f2864])).
fof(f6343, plain, (~ spl130_112 | ~ spl130_122), inference(avatar_split_clause, [], [f6329, f1709, f1667])).
fof(f6329, plain, (~ (e1 = op(e0, e2)) | ~ spl130_122), inference(backward_demodulation, [], [f645, f1711])).
fof(f6342, plain, (~ spl130_72 | ~ spl130_122), inference(avatar_split_clause, [], [f6327, f1709, f1499])).
fof(f6327, plain, (~ (e1 = op(e2, e0)) | ~ spl130_122), inference(backward_demodulation, [], [f595, f1711])).
fof(f6293, plain, (~ spl130_74 | ~ spl130_64), inference(avatar_split_clause, [], [f6292, f1465, f1507])).
fof(f6292, plain, (~ (e3 = op(e2, e0)) | ~ spl130_64), inference(forward_demodulation, [], [f665, f1467])).
fof(f6282, plain, (~ spl130_61 | ~ spl130_64), inference(avatar_contradiction_clause, [], [f6281])).
fof(f6281, plain, ($false | (~ spl130_61 | ~ spl130_64)), inference(subsumption_resolution, [], [f6280, f696])).
fof(f696, plain, ~ (e0 = e3), inference(cnf_transformation, [], [f5])).
fof(f6280, plain, ((e0 = e3) | (~ spl130_61 | ~ spl130_64)), inference(forward_demodulation, [], [f1467, f1455])).
fof(f6276, plain, (~ spl130_76 | ~ spl130_78), inference(avatar_contradiction_clause, [], [f6275])).
fof(f6275, plain, ($false | (~ spl130_76 | ~ spl130_78)), inference(subsumption_resolution, [], [f6274, f695])).
fof(f6274, plain, ((e0 = e2) | (~ spl130_76 | ~ spl130_78)), inference(forward_demodulation, [], [f1526, f1518])).
fof(f6209, plain, (~ spl130_125 | ~ spl130_61 | spl130_263), inference(avatar_split_clause, [], [f6118, f2676, f1453, f1721])).
fof(f2676, plain, (spl130_263 <=> (e4 = op(op(e2, e2), op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl130_263])])).
fof(f6118, plain, (~ (op(e0, e0) = e4) | (~ spl130_61 | spl130_263)), inference(backward_demodulation, [], [f2678, f1455])).
fof(f2678, plain, (~ (e4 = op(op(e2, e2), op(e2, e2))) | spl130_263), inference(avatar_component_clause, [], [f2676])).
fof(f6208, plain, (~ spl130_124 | ~ spl130_61 | spl130_264), inference(avatar_split_clause, [], [f6119, f2681, f1453, f1717])).
fof(f2681, plain, (spl130_264 <=> (e3 = op(op(e2, e2), op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl130_264])])).
fof(f6119, plain, (~ (op(e0, e0) = e3) | (~ spl130_61 | spl130_264)), inference(backward_demodulation, [], [f2683, f1455])).
fof(f2683, plain, (~ (e3 = op(op(e2, e2), op(e2, e2))) | spl130_264), inference(avatar_component_clause, [], [f2681])).
fof(f6197, plain, (~ spl130_112 | ~ spl130_12), inference(avatar_split_clause, [], [f6196, f1247, f1667])).
fof(f6196, plain, (~ (e1 = op(e0, e2)) | ~ spl130_12), inference(forward_demodulation, [], [f617, f1249])).
fof(f617, plain, ~ (op(e0, e2) = op(e4, e2)), inference(cnf_transformation, [], [f4])).
fof(f6193, plain, (~ spl130_102 | ~ spl130_52), inference(avatar_split_clause, [], [f6192, f1415, f1625])).
fof(f6192, plain, (~ (e1 = op(e0, e4)) | ~ spl130_52), inference(forward_demodulation, [], [f635, f1417])).
fof(f1417, plain, ((e1 = op(e2, e4)) | ~ spl130_52), inference(avatar_component_clause, [], [f1415])).
fof(f6172, plain, (~ spl130_72 | ~ spl130_52), inference(avatar_split_clause, [], [f6171, f1415, f1499])).
fof(f6171, plain, (~ (e1 = op(e2, e0)) | ~ spl130_52), inference(forward_demodulation, [], [f667, f1417])).
fof(f6164, plain, (~ spl130_93 | ~ spl130_122 | spl130_302), inference(avatar_contradiction_clause, [], [f6163])).
fof(f6163, plain, ($false | (~ spl130_93 | ~ spl130_122 | spl130_302)), inference(subsumption_resolution, [], [f6162, f1589])).
fof(f6162, plain, (~ (e2 = op(e1, e1)) | (~ spl130_122 | spl130_302)), inference(forward_demodulation, [], [f2871, f1711])).
fof(f6107, plain, (~ spl130_70 | ~ spl130_93 | spl130_286), inference(avatar_contradiction_clause, [], [f6106])).
fof(f6106, plain, ($false | (~ spl130_70 | ~ spl130_93 | spl130_286)), inference(subsumption_resolution, [], [f6099, f1492])).
fof(f6099, plain, (~ (e4 = op(e2, e1)) | (~ spl130_93 | spl130_286)), inference(backward_demodulation, [], [f2788, f1589])).
fof(f6089, plain, (~ spl130_4 | ~ spl130_100 | ~ spl130_122 | spl130_316), inference(avatar_contradiction_clause, [], [f6088])).
fof(f6088, plain, ($false | (~ spl130_4 | ~ spl130_100 | ~ spl130_122 | spl130_316)), inference(subsumption_resolution, [], [f6086, f1215])).
fof(f6086, plain, (~ (e3 = op(e4, e4)) | (~ spl130_100 | ~ spl130_122 | spl130_316)), inference(backward_demodulation, [], [f6078, f1618])).
fof(f6078, plain, (~ (e3 = op(op(e1, e0), op(e1, e0))) | (~ spl130_122 | spl130_316)), inference(backward_demodulation, [], [f2968, f1711])).
fof(f6081, plain, (~ spl130_98 | ~ spl130_122 | spl130_296), inference(avatar_split_clause, [], [f6073, f2838, f1709, f1608])).
fof(f6073, plain, (~ (e2 = op(e1, e0)) | (~ spl130_122 | spl130_296)), inference(backward_demodulation, [], [f2840, f1711])).
fof(f6041, plain, (~ spl130_26 | ~ spl130_76), inference(avatar_split_clause, [], [f6040, f1516, f1306])).
fof(f6040, plain, (~ (e0 = op(e3, e4)) | ~ spl130_76), inference(forward_demodulation, [], [f639, f1518])).
fof(f6031, plain, (~ spl130_63 | ~ spl130_58), inference(avatar_split_clause, [], [f6030, f1440, f1461])).
fof(f6030, plain, (~ (e2 = op(e2, e2)) | ~ spl130_58), inference(forward_demodulation, [], [f671, f1442])).
fof(f6012, plain, (~ spl130_27 | ~ spl130_42), inference(avatar_split_clause, [], [f6011, f1373, f1310])).
fof(f6011, plain, (~ (e1 = op(e3, e4)) | ~ spl130_42), inference(forward_demodulation, [], [f680, f1375])).
fof(f6006, plain, (~ spl130_15 | ~ spl130_10), inference(avatar_split_clause, [], [f6005, f1238, f1259])).
fof(f6005, plain, (~ (e4 = op(e4, e2)) | ~ spl130_10), inference(forward_demodulation, [], [f691, f1240])).
fof(f5939, plain, (~ spl130_66 | ~ spl130_70), inference(avatar_contradiction_clause, [], [f5938])).
fof(f5938, plain, ($false | (~ spl130_66 | ~ spl130_70)), inference(subsumption_resolution, [], [f5937, f697])).
fof(f5937, plain, ((e0 = e4) | (~ spl130_66 | ~ spl130_70)), inference(backward_demodulation, [], [f1492, f1476])).
fof(f5927, plain, (~ spl130_76 | ~ spl130_80), inference(avatar_contradiction_clause, [], [f5926])).
fof(f5926, plain, ($false | (~ spl130_76 | ~ spl130_80)), inference(subsumption_resolution, [], [f5925, f697])).
fof(f5925, plain, ((e0 = e4) | (~ spl130_76 | ~ spl130_80)), inference(backward_demodulation, [], [f1534, f1518])).
fof(f1534, plain, ((e4 = op(e1, e4)) | ~ spl130_80), inference(avatar_component_clause, [], [f1532])).
fof(f5909, plain, (~ spl130_83 | ~ spl130_98), inference(avatar_split_clause, [], [f5904, f1608, f1545])).
fof(f5904, plain, (~ (e2 = op(e1, e3)) | ~ spl130_98), inference(backward_demodulation, [], [f656, f1610])).
fof(f5908, plain, (~ spl130_88 | ~ spl130_98), inference(avatar_split_clause, [], [f5903, f1608, f1566])).
fof(f5903, plain, (~ (e2 = op(e1, e2)) | ~ spl130_98), inference(backward_demodulation, [], [f655, f1610])).
fof(f5907, plain, (~ spl130_23 | ~ spl130_98), inference(avatar_split_clause, [], [f5902, f1608, f1293])).
fof(f5902, plain, (~ (e2 = op(e4, e0)) | ~ spl130_98), inference(backward_demodulation, [], [f600, f1610])).
fof(f5887, plain, (~ spl130_110 | ~ spl130_115), inference(avatar_split_clause, [], [f5885, f1679, f1658])).
fof(f5885, plain, (~ (e4 = op(e0, e3)) | ~ spl130_115), inference(backward_demodulation, [], [f651, f1681])).
fof(f5880, plain, (~ spl130_44 | ~ spl130_119), inference(avatar_split_clause, [], [f5875, f1696, f1381])).
fof(f5875, plain, (~ (e3 = op(e3, e1)) | ~ spl130_119), inference(backward_demodulation, [], [f606, f1698])).
fof(f5866, plain, (~ spl130_50 | ~ spl130_124 | spl130_304), inference(avatar_contradiction_clause, [], [f5865])).
fof(f5865, plain, ($false | (~ spl130_50 | ~ spl130_124 | spl130_304)), inference(subsumption_resolution, [], [f5858, f1408])).
fof(f5858, plain, (~ (e4 = op(e3, e0)) | (~ spl130_124 | spl130_304)), inference(backward_demodulation, [], [f2884, f1719])).
fof(f5848, plain, (~ spl130_26 | ~ spl130_129), inference(avatar_contradiction_clause, [], [f5847])).
fof(f5847, plain, ($false | (~ spl130_26 | ~ spl130_129)), inference(subsumption_resolution, [], [f5846, f697])).
fof(f5846, plain, ((e0 = e4) | (~ spl130_26 | ~ spl130_129)), inference(forward_demodulation, [], [f5835, f1308])).
fof(f1308, plain, ((e0 = op(e3, e4)) | ~ spl130_26), inference(avatar_component_clause, [], [f1306])).
fof(f5835, plain, ((e4 = op(e3, e4)) | ~ spl130_129), inference(backward_demodulation, [], [f541, f1740])).
fof(f5826, plain, (~ spl130_33 | ~ spl130_4 | spl130_272), inference(avatar_split_clause, [], [f5825, f2717, f1213, f1335])).
fof(f5825, plain, (~ (e2 = op(e3, e3)) | (~ spl130_4 | spl130_272)), inference(forward_demodulation, [], [f2719, f1215])).
fof(f5807, plain, (~ spl130_99 | ~ spl130_74), inference(avatar_split_clause, [], [f5729, f1507, f1612])).
fof(f5729, plain, (~ (e3 = op(e1, e0)) | ~ spl130_74), inference(forward_demodulation, [], [f598, f1509])).
fof(f5776, plain, (~ spl130_4 | ~ spl130_65 | spl130_264), inference(avatar_contradiction_clause, [], [f5775])).
fof(f5775, plain, ($false | (~ spl130_4 | ~ spl130_65 | spl130_264)), inference(subsumption_resolution, [], [f5774, f1215])).
fof(f5774, plain, (~ (e3 = op(e4, e4)) | (~ spl130_65 | spl130_264)), inference(forward_demodulation, [], [f2683, f1471])).
fof(f5744, plain, (~ spl130_4 | ~ spl130_26 | ~ spl130_122 | spl130_293), inference(avatar_contradiction_clause, [], [f5743])).
fof(f5743, plain, ($false | (~ spl130_4 | ~ spl130_26 | ~ spl130_122 | spl130_293)), inference(subsumption_resolution, [], [f5742, f1711])).
fof(f5742, plain, (~ (op(e0, e0) = e1) | (~ spl130_4 | ~ spl130_26 | spl130_293)), inference(forward_demodulation, [], [f5741, f1308])).
fof(f5741, plain, (~ (e1 = op(op(e3, e4), op(e3, e4))) | (~ spl130_4 | spl130_293)), inference(forward_demodulation, [], [f2826, f1215])).
fof(f5681, plain, (~ spl130_15 | ~ spl130_20), inference(avatar_split_clause, [], [f5679, f1280, f1259])).
fof(f5679, plain, (~ (e4 = op(e4, e2)) | ~ spl130_20), inference(backward_demodulation, [], [f688, f1282])).
fof(f1282, plain, ((e4 = op(e4, e1)) | ~ spl130_20), inference(avatar_component_clause, [], [f1280])).
fof(f5678, plain, (~ spl130_3 | ~ spl130_23), inference(avatar_split_clause, [], [f5675, f1293, f1209])).
fof(f5675, plain, (~ (e2 = op(e4, e4)) | ~ spl130_23), inference(backward_demodulation, [], [f687, f1295])).
fof(f5667, plain, (~ spl130_33 | ~ spl130_56 | spl130_292), inference(avatar_contradiction_clause, [], [f5666])).
fof(f5666, plain, ($false | (~ spl130_33 | ~ spl130_56 | spl130_292)), inference(subsumption_resolution, [], [f5662, f1434])).
fof(f5662, plain, (~ (e0 = op(e2, e3)) | (~ spl130_33 | spl130_292)), inference(backward_demodulation, [], [f2820, f1337])).
fof(f2820, plain, (~ (e0 = op(op(e3, e3), e3)) | spl130_292), inference(avatar_component_clause, [], [f2818])).
fof(f5611, plain, (~ spl130_58 | ~ spl130_68), inference(avatar_split_clause, [], [f5605, f1482, f1440])).
fof(f5605, plain, (~ (e2 = op(e2, e3)) | ~ spl130_68), inference(backward_demodulation, [], [f669, f1484])).
fof(f5610, plain, (~ spl130_63 | ~ spl130_68), inference(avatar_split_clause, [], [f5604, f1482, f1461])).
fof(f5604, plain, (~ (e2 = op(e2, e2)) | ~ spl130_68), inference(backward_demodulation, [], [f668, f1484])).
fof(f5601, plain, (~ spl130_59 | ~ spl130_74), inference(avatar_split_clause, [], [f5598, f1507, f1444])).
fof(f5598, plain, (~ (e3 = op(e2, e3)) | ~ spl130_74), inference(backward_demodulation, [], [f666, f1509])).
fof(f5593, plain, (~ spl130_55 | ~ spl130_80), inference(avatar_split_clause, [], [f5590, f1532, f1427])).
fof(f5590, plain, (~ (e4 = op(e2, e4)) | ~ spl130_80), inference(backward_demodulation, [], [f638, f1534])).
fof(f5589, plain, (~ spl130_63 | ~ spl130_88), inference(avatar_split_clause, [], [f5583, f1566, f1461])).
fof(f5583, plain, (~ (e2 = op(e2, e2)) | ~ spl130_88), inference(backward_demodulation, [], [f618, f1568])).
fof(f5579, plain, (~ spl130_96 | ~ spl130_100), inference(avatar_contradiction_clause, [], [f5578])).
fof(f5578, plain, ($false | (~ spl130_96 | ~ spl130_100)), inference(subsumption_resolution, [], [f5577, f697])).
fof(f5577, plain, ((e0 = e4) | (~ spl130_96 | ~ spl130_100)), inference(backward_demodulation, [], [f1618, f1602])).
fof(f5573, plain, (~ spl130_3 | ~ spl130_103), inference(avatar_split_clause, [], [f5570, f1629, f1209])).
fof(f5570, plain, (~ (e2 = op(e4, e4)) | ~ spl130_103), inference(backward_demodulation, [], [f637, f1631])).
fof(f637, plain, ~ (op(e0, e4) = op(e4, e4)), inference(cnf_transformation, [], [f4])).
fof(f5565, plain, (~ spl130_39 | ~ spl130_114), inference(avatar_split_clause, [], [f5561, f1675, f1360])).
fof(f5561, plain, (~ (e3 = op(e3, e2)) | ~ spl130_114), inference(backward_demodulation, [], [f616, f1677])).
fof(f616, plain, ~ (op(e0, e2) = op(e3, e2)), inference(cnf_transformation, [], [f4])).
fof(f5502, plain, (~ spl130_124 | ~ spl130_290 | spl130_314), inference(avatar_split_clause, [], [f5498, f2954, f2808, f1717])).
fof(f2954, plain, (spl130_314 <=> (e3 = op(op(op(e2, e2), e2), op(op(e2, e2), e2)))), introduced(avatar_definition, [new_symbols(naming, [spl130_314])])).
fof(f5498, plain, (~ (op(e0, e0) = e3) | (~ spl130_290 | spl130_314)), inference(backward_demodulation, [], [f2956, f2809])).
fof(f2956, plain, (~ (e3 = op(op(op(e2, e2), e2), op(op(e2, e2), e2))) | spl130_314), inference(avatar_component_clause, [], [f2954])).
fof(f5479, plain, (~ spl130_80 | ~ spl130_100), inference(avatar_split_clause, [], [f5398, f1616, f1532])).
fof(f5398, plain, (~ (e4 = op(e1, e4)) | ~ spl130_100), inference(backward_demodulation, [], [f657, f1618])).
fof(f5469, plain, (~ spl130_59 | ~ spl130_84), inference(avatar_split_clause, [], [f5468, f1549, f1444])).
fof(f5468, plain, (~ (e3 = op(e2, e3)) | ~ spl130_84), inference(forward_demodulation, [], [f628, f1551])).
fof(f5467, plain, (~ spl130_60 | ~ spl130_110), inference(avatar_split_clause, [], [f5466, f1658, f1448])).
fof(f5466, plain, (~ (e4 = op(e2, e3)) | ~ spl130_110), inference(forward_demodulation, [], [f625, f1660])).
fof(f5459, plain, (~ spl130_31 | spl130_273 | ~ spl130_281), inference(avatar_split_clause, [], [f5352, f2760, f2722, f1327])).
fof(f2760, plain, (spl130_281 <=> (e3 = op(op(e1, e1), e1))), introduced(avatar_definition, [new_symbols(naming, [spl130_281])])).
fof(f5352, plain, (~ (e0 = op(e3, e3)) | (spl130_273 | ~ spl130_281)), inference(forward_demodulation, [], [f2724, f2761])).
fof(f2761, plain, ((e3 = op(op(e1, e1), e1)) | ~ spl130_281), inference(avatar_component_clause, [], [f2760])).
fof(f5446, plain, (~ spl130_16 | ~ spl130_19), inference(avatar_contradiction_clause, [], [f5445])).
fof(f5445, plain, ($false | (~ spl130_16 | ~ spl130_19)), inference(subsumption_resolution, [], [f5444, f696])).
fof(f5444, plain, ((e0 = e3) | (~ spl130_16 | ~ spl130_19)), inference(backward_demodulation, [], [f1278, f1266])).
fof(f5399, plain, (~ spl130_25 | ~ spl130_100), inference(avatar_split_clause, [], [f5395, f1616, f1301])).
fof(f5395, plain, (~ (e4 = op(e4, e0)) | ~ spl130_100), inference(backward_demodulation, [], [f600, f1618])).
fof(f5393, plain, (~ spl130_102 | ~ spl130_104), inference(avatar_contradiction_clause, [], [f5392])).
fof(f5392, plain, ($false | (~ spl130_102 | ~ spl130_104)), inference(subsumption_resolution, [], [f5391, f699])).
fof(f5391, plain, ((e1 = e3) | (~ spl130_102 | ~ spl130_104)), inference(backward_demodulation, [], [f1635, f1627])).
fof(f5381, plain, (~ spl130_103 | ~ spl130_118), inference(avatar_split_clause, [], [f5376, f1692, f1629])).
fof(f5376, plain, (~ (e2 = op(e0, e4)) | ~ spl130_118), inference(backward_demodulation, [], [f650, f1694])).
fof(f5372, plain, (spl130_49 | ~ spl130_124 | ~ spl130_301), inference(avatar_split_clause, [], [f5365, f2864, f1717, f1402])).
fof(f5365, plain, ((e3 = op(e3, e0)) | (~ spl130_124 | ~ spl130_301)), inference(backward_demodulation, [], [f2865, f1719])).
fof(f5370, plain, (~ spl130_48 | ~ spl130_124 | spl130_296), inference(avatar_split_clause, [], [f5363, f2838, f1717, f1398])).
fof(f5363, plain, (~ (e2 = op(e3, e0)) | (~ spl130_124 | spl130_296)), inference(backward_demodulation, [], [f2840, f1719])).
fof(f5369, plain, (~ spl130_104 | ~ spl130_124), inference(avatar_split_clause, [], [f5361, f1717, f1633])).
fof(f5361, plain, (~ (e3 = op(e0, e4)) | ~ spl130_124), inference(backward_demodulation, [], [f647, f1719])).
fof(f5308, plain, (~ spl130_108 | ~ spl130_31 | spl130_277), inference(avatar_split_clause, [], [f5307, f2740, f1327, f1650])).
fof(f5307, plain, (~ (e2 = op(e0, e3)) | (~ spl130_31 | spl130_277)), inference(forward_demodulation, [], [f2742, f1329])).
fof(f2742, plain, (~ (e2 = op(op(e3, e3), e3)) | spl130_277), inference(avatar_component_clause, [], [f2740])).
fof(f5285, plain, (~ spl130_23 | ~ spl130_25), inference(avatar_contradiction_clause, [], [f5284])).
fof(f5284, plain, ($false | (~ spl130_23 | ~ spl130_25)), inference(subsumption_resolution, [], [f5283, f702])).
fof(f702, plain, ~ (e2 = e4), inference(cnf_transformation, [], [f5])).
fof(f5283, plain, ((e2 = e4) | (~ spl130_23 | ~ spl130_25)), inference(backward_demodulation, [], [f1303, f1295])).
fof(f5278, plain, (~ spl130_2 | ~ spl130_27), inference(avatar_split_clause, [], [f5277, f1310, f1205])).
fof(f5277, plain, (~ (e1 = op(e4, e4)) | ~ spl130_27), inference(backward_demodulation, [], [f643, f1312])).
fof(f1312, plain, ((e1 = op(e3, e4)) | ~ spl130_27), inference(avatar_component_clause, [], [f1310])).
fof(f5273, plain, (~ spl130_31 | ~ spl130_33), inference(avatar_contradiction_clause, [], [f5272])).
fof(f5272, plain, ($false | (~ spl130_31 | ~ spl130_33)), inference(subsumption_resolution, [], [f5271, f695])).
fof(f5271, plain, ((e0 = e2) | (~ spl130_31 | ~ spl130_33)), inference(backward_demodulation, [], [f1337, f1329])).
fof(f5269, plain, (~ spl130_43 | ~ spl130_45), inference(avatar_contradiction_clause, [], [f5268])).
fof(f5268, plain, ($false | (~ spl130_43 | ~ spl130_45)), inference(subsumption_resolution, [], [f5266, f702])).
fof(f5266, plain, ((e2 = e4) | (~ spl130_43 | ~ spl130_45)), inference(backward_demodulation, [], [f1387, f1379])).
fof(f5245, plain, (~ spl130_107 | ~ spl130_122), inference(avatar_split_clause, [], [f5235, f1709, f1646])).
fof(f5235, plain, (~ (e1 = op(e0, e3)) | ~ spl130_122), inference(backward_demodulation, [], [f646, f1711])).
fof(f5224, plain, (spl130_120 | ~ spl130_91 | ~ spl130_286), inference(avatar_split_clause, [], [f5223, f2786, f1579, f1700])).
fof(f5223, plain, ((e4 = op(e0, e1)) | (~ spl130_91 | ~ spl130_286)), inference(forward_demodulation, [], [f2787, f1581])).
fof(f5168, plain, (~ spl130_2 | ~ spl130_125 | spl130_311), inference(avatar_contradiction_clause, [], [f5167])).
fof(f5167, plain, ($false | (~ spl130_2 | ~ spl130_125 | spl130_311)), inference(subsumption_resolution, [], [f5166, f1207])).
fof(f5166, plain, (~ (e1 = op(e4, e4)) | (~ spl130_125 | spl130_311)), inference(forward_demodulation, [], [f2931, f1723])).
fof(f5136, plain, (~ spl130_24 | ~ spl130_125 | spl130_301), inference(avatar_contradiction_clause, [], [f5135])).
fof(f5135, plain, ($false | (~ spl130_24 | ~ spl130_125 | spl130_301)), inference(subsumption_resolution, [], [f5134, f1299])).
fof(f5134, plain, (~ (e3 = op(e4, e0)) | (~ spl130_125 | spl130_301)), inference(forward_demodulation, [], [f2866, f1723])).
fof(f5093, plain, (~ spl130_12 | ~ spl130_15), inference(avatar_contradiction_clause, [], [f5092])).
fof(f5092, plain, ($false | (~ spl130_12 | ~ spl130_15)), inference(subsumption_resolution, [], [f5091, f700])).
fof(f5091, plain, ((e1 = e4) | (~ spl130_12 | ~ spl130_15)), inference(forward_demodulation, [], [f1261, f1249])).
fof(f1261, plain, ((e4 = op(e4, e2)) | ~ spl130_15), inference(avatar_component_clause, [], [f1259])).
fof(f5090, plain, (~ spl130_16 | ~ spl130_18), inference(avatar_contradiction_clause, [], [f5089])).
fof(f5089, plain, ($false | (~ spl130_16 | ~ spl130_18)), inference(subsumption_resolution, [], [f5088, f695])).
fof(f5088, plain, ((e0 = e2) | (~ spl130_16 | ~ spl130_18)), inference(backward_demodulation, [], [f1274, f1266])).
fof(f5068, plain, (~ spl130_52 | ~ spl130_55), inference(avatar_contradiction_clause, [], [f5067])).
fof(f5067, plain, ($false | (~ spl130_52 | ~ spl130_55)), inference(subsumption_resolution, [], [f5066, f700])).
fof(f5066, plain, ((e1 = e4) | (~ spl130_52 | ~ spl130_55)), inference(backward_demodulation, [], [f1429, f1417])).
fof(f1429, plain, ((e4 = op(e2, e4)) | ~ spl130_55), inference(avatar_component_clause, [], [f1427])).
fof(f5049, plain, (~ spl130_76 | ~ spl130_79), inference(avatar_contradiction_clause, [], [f5048])).
fof(f5048, plain, ($false | (~ spl130_76 | ~ spl130_79)), inference(subsumption_resolution, [], [f5047, f696])).
fof(f5047, plain, ((e0 = e3) | (~ spl130_76 | ~ spl130_79)), inference(backward_demodulation, [], [f1530, f1518])).
fof(f4946, plain, (~ spl130_12 | ~ spl130_128), inference(avatar_contradiction_clause, [], [f4945])).
fof(f4945, plain, ($false | (~ spl130_12 | ~ spl130_128)), inference(subsumption_resolution, [], [f4944, f700])).
fof(f4944, plain, ((e1 = e4) | (~ spl130_12 | ~ spl130_128)), inference(forward_demodulation, [], [f4747, f1249])).
fof(f4747, plain, ((e4 = op(e4, e2)) | ~ spl130_128), inference(backward_demodulation, [], [f542, f1736])).
fof(f4914, plain, (~ spl130_1 | ~ spl130_4), inference(avatar_contradiction_clause, [], [f4913])).
fof(f4913, plain, ($false | (~ spl130_1 | ~ spl130_4)), inference(subsumption_resolution, [], [f4912, f696])).
fof(f4912, plain, ((e0 = e3) | (~ spl130_1 | ~ spl130_4)), inference(backward_demodulation, [], [f1215, f1203])).
fof(f4872, plain, (~ spl130_41 | ~ spl130_45), inference(avatar_contradiction_clause, [], [f4871])).
fof(f4871, plain, ($false | (~ spl130_41 | ~ spl130_45)), inference(subsumption_resolution, [], [f4869, f697])).
fof(f4869, plain, ((e0 = e4) | (~ spl130_41 | ~ spl130_45)), inference(backward_demodulation, [], [f1387, f1371])).
fof(f4866, plain, (~ spl130_30 | ~ spl130_45), inference(avatar_split_clause, [], [f4865, f1385, f1322])).
fof(f4865, plain, (~ (e4 = op(e3, e4)) | ~ spl130_45), inference(backward_demodulation, [], [f680, f1387])).
fof(f4862, plain, (~ spl130_27 | ~ spl130_47), inference(avatar_split_clause, [], [f4856, f1394, f1310])).
fof(f4856, plain, (~ (e1 = op(e3, e4)) | ~ spl130_47), inference(backward_demodulation, [], [f677, f1396])).
fof(f4860, plain, (~ spl130_22 | ~ spl130_47), inference(avatar_split_clause, [], [f4853, f1394, f1289])).
fof(f4853, plain, (~ (e1 = op(e4, e0)) | ~ spl130_47), inference(backward_demodulation, [], [f603, f1396])).
fof(f4829, plain, (~ spl130_71 | ~ spl130_73), inference(avatar_contradiction_clause, [], [f4828])).
fof(f4828, plain, ($false | (~ spl130_71 | ~ spl130_73)), inference(subsumption_resolution, [], [f4827, f695])).
fof(f4827, plain, ((e0 = e2) | (~ spl130_71 | ~ spl130_73)), inference(backward_demodulation, [], [f1505, f1497])).
fof(f4825, plain, (~ spl130_53 | ~ spl130_73), inference(avatar_split_clause, [], [f4822, f1503, f1419])).
fof(f4822, plain, (~ (e2 = op(e2, e4)) | ~ spl130_73), inference(backward_demodulation, [], [f667, f1505])).
fof(f4800, plain, (~ spl130_93 | ~ spl130_95), inference(avatar_contradiction_clause, [], [f4799])).
fof(f4799, plain, ($false | (~ spl130_93 | ~ spl130_95)), inference(subsumption_resolution, [], [f4798, f702])).
fof(f4798, plain, ((e2 = e4) | (~ spl130_93 | ~ spl130_95)), inference(backward_demodulation, [], [f1597, f1589])).
fof(f4788, plain, (~ spl130_53 | ~ spl130_103), inference(avatar_split_clause, [], [f4783, f1629, f1419])).
fof(f4783, plain, (~ (e2 = op(e2, e4)) | ~ spl130_103), inference(backward_demodulation, [], [f635, f1631])).
fof(f4765, plain, (~ spl130_124 | ~ spl130_125), inference(avatar_contradiction_clause, [], [f4764])).
fof(f4764, plain, ($false | (~ spl130_124 | ~ spl130_125)), inference(subsumption_resolution, [], [f4763, f703])).
fof(f703, plain, ~ (e3 = e4), inference(cnf_transformation, [], [f5])).
fof(f4763, plain, ((e3 = e4) | (~ spl130_124 | ~ spl130_125)), inference(backward_demodulation, [], [f1723, f1719])).
fof(f4732, plain, (~ spl130_22 | ~ spl130_125 | spl130_310), inference(avatar_split_clause, [], [f4731, f2918, f1721, f1289])).
fof(f4731, plain, (~ (e1 = op(e4, e0)) | (~ spl130_125 | spl130_310)), inference(forward_demodulation, [], [f2920, f1723])).
fof(f4725, plain, (~ spl130_69 | ~ spl130_119), inference(avatar_contradiction_clause, [], [f4724])).
fof(f4724, plain, ($false | (~ spl130_69 | ~ spl130_119)), inference(subsumption_resolution, [], [f4723, f1698])).
fof(f4723, plain, (~ (e3 = op(e0, e1)) | ~ spl130_69), inference(forward_demodulation, [], [f605, f1488])).
fof(f4685, plain, (~ spl130_25 | ~ spl130_125), inference(avatar_split_clause, [], [f4197, f1721, f1301])).
fof(f4197, plain, (~ (e4 = op(e4, e0)) | ~ spl130_125), inference(backward_demodulation, [], [f597, f1723])).
fof(f4646, plain, (~ spl130_21 | ~ spl130_22), inference(avatar_contradiction_clause, [], [f4645])).
fof(f4645, plain, ($false | (~ spl130_21 | ~ spl130_22)), inference(subsumption_resolution, [], [f4644, f694])).
fof(f4644, plain, ((e0 = e1) | (~ spl130_21 | ~ spl130_22)), inference(forward_demodulation, [], [f1291, f1287])).
fof(f4643, plain, (~ spl130_30 | ~ spl130_40), inference(avatar_split_clause, [], [f4642, f1364, f1322])).
fof(f4642, plain, (~ (e4 = op(e3, e4)) | ~ spl130_40), inference(backward_demodulation, [], [f682, f1366])).
fof(f1366, plain, ((e4 = op(e3, e2)) | ~ spl130_40), inference(avatar_component_clause, [], [f1364])).
fof(f4622, plain, (~ spl130_54 | ~ spl130_69), inference(avatar_split_clause, [], [f4619, f1486, f1423])).
fof(f4619, plain, (~ (e3 = op(e2, e4)) | ~ spl130_69), inference(backward_demodulation, [], [f670, f1488])).
fof(f4616, plain, (~ spl130_74 | ~ spl130_75), inference(avatar_contradiction_clause, [], [f4615])).
fof(f4615, plain, ($false | (~ spl130_74 | ~ spl130_75)), inference(subsumption_resolution, [], [f4614, f703])).
fof(f4614, plain, ((e3 = e4) | (~ spl130_74 | ~ spl130_75)), inference(backward_demodulation, [], [f1513, f1509])).
fof(f4612, plain, (~ spl130_50 | ~ spl130_75), inference(avatar_split_clause, [], [f4608, f1511, f1406])).
fof(f4608, plain, (~ (e4 = op(e3, e0)) | ~ spl130_75), inference(backward_demodulation, [], [f601, f1513])).
fof(f4607, plain, (~ spl130_60 | ~ spl130_85), inference(avatar_split_clause, [], [f4606, f1553, f1448])).
fof(f4606, plain, (~ (e4 = op(e2, e3)) | ~ spl130_85), inference(backward_demodulation, [], [f628, f1555])).
fof(f4605, plain, (~ spl130_86 | ~ spl130_89), inference(avatar_contradiction_clause, [], [f4604])).
fof(f4604, plain, ($false | (~ spl130_86 | ~ spl130_89)), inference(subsumption_resolution, [], [f4601, f696])).
fof(f4601, plain, ((e0 = e3) | (~ spl130_86 | ~ spl130_89)), inference(backward_demodulation, [], [f1572, f1560])).
fof(f4516, plain, (~ spl130_39 | ~ spl130_34), inference(avatar_split_clause, [], [f4515, f1339, f1360])).
fof(f4515, plain, (~ (e3 = op(e3, e2)) | ~ spl130_34), inference(forward_demodulation, [], [f681, f1341])).
fof(f4496, plain, (~ spl130_16 | ~ spl130_95 | spl130_306), inference(avatar_split_clause, [], [f4409, f2896, f1595, f1264])).
fof(f4409, plain, (~ (e0 = op(e4, e1)) | (~ spl130_95 | spl130_306)), inference(backward_demodulation, [], [f2898, f1597])).
fof(f4494, plain, (~ spl130_7 | ~ spl130_130), inference(avatar_contradiction_clause, [], [f4493])).
fof(f4493, plain, ($false | (~ spl130_7 | ~ spl130_130)), inference(subsumption_resolution, [], [f4492, f699])).
fof(f4492, plain, ((e1 = e3) | (~ spl130_7 | ~ spl130_130)), inference(forward_demodulation, [], [f4370, f1228])).
fof(f1228, plain, ((e1 = op(e4, e3)) | ~ spl130_7), inference(avatar_component_clause, [], [f1226])).
fof(f4370, plain, ((e3 = op(e4, e3)) | ~ spl130_130), inference(backward_demodulation, [], [f539, f1744])).
fof(f4437, plain, (~ spl130_72 | ~ spl130_73), inference(avatar_contradiction_clause, [], [f4436])).
fof(f4436, plain, ($false | (~ spl130_72 | ~ spl130_73)), inference(subsumption_resolution, [], [f4435, f698])).
fof(f698, plain, ~ (e1 = e2), inference(cnf_transformation, [], [f5])).
fof(f4435, plain, ((e1 = e2) | (~ spl130_72 | ~ spl130_73)), inference(backward_demodulation, [], [f1505, f1501])).
fof(f4424, plain, (~ spl130_81 | ~ spl130_82), inference(avatar_contradiction_clause, [], [f4423])).
fof(f4423, plain, ($false | (~ spl130_81 | ~ spl130_82)), inference(subsumption_resolution, [], [f4422, f694])).
fof(f4422, plain, ((e0 = e1) | (~ spl130_81 | ~ spl130_82)), inference(backward_demodulation, [], [f1543, f1539])).
fof(f4413, plain, (~ spl130_19 | ~ spl130_95 | spl130_281), inference(avatar_split_clause, [], [f4407, f2760, f1595, f1276])).
fof(f4407, plain, (~ (e3 = op(e4, e1)) | (~ spl130_95 | spl130_281)), inference(backward_demodulation, [], [f2762, f1597])).
fof(f2762, plain, (~ (e3 = op(op(e1, e1), e1)) | spl130_281), inference(avatar_component_clause, [], [f2760])).
fof(f4401, plain, (~ spl130_89 | ~ spl130_99), inference(avatar_split_clause, [], [f4395, f1612, f1570])).
fof(f4395, plain, (~ (e3 = op(e1, e2)) | ~ spl130_99), inference(backward_demodulation, [], [f655, f1614])).
fof(f4399, plain, (~ spl130_49 | ~ spl130_99), inference(avatar_split_clause, [], [f4392, f1612, f1402])).
fof(f4392, plain, (~ (e3 = op(e3, e0)) | ~ spl130_99), inference(backward_demodulation, [], [f599, f1614])).
fof(f4385, plain, (~ spl130_28 | ~ spl130_130), inference(avatar_contradiction_clause, [], [f4384])).
fof(f4384, plain, ($false | (~ spl130_28 | ~ spl130_130)), inference(subsumption_resolution, [], [f4383, f701])).
fof(f701, plain, ~ (e2 = e3), inference(cnf_transformation, [], [f5])).
fof(f4383, plain, ((e2 = e3) | (~ spl130_28 | ~ spl130_130)), inference(forward_demodulation, [], [f4371, f1316])).
fof(f4371, plain, ((e3 = op(e3, e4)) | ~ spl130_130), inference(backward_demodulation, [], [f540, f1744])).
fof(f4358, plain, (~ spl130_89 | ~ spl130_62 | spl130_283), inference(avatar_split_clause, [], [f4033, f2770, f1457, f1570])).
fof(f4033, plain, (~ (e3 = op(e1, e2)) | (~ spl130_62 | spl130_283)), inference(backward_demodulation, [], [f2772, f1459])).
fof(f2772, plain, (~ (e3 = op(op(e2, e2), e2)) | spl130_283), inference(avatar_component_clause, [], [f2770])).
fof(f4349, plain, (~ spl130_97 | ~ spl130_82), inference(avatar_split_clause, [], [f4348, f1541, f1604])).
fof(f4348, plain, (~ (e1 = op(e1, e0)) | ~ spl130_82), inference(forward_demodulation, [], [f656, f1543])).
fof(f4329, plain, (~ spl130_50 | ~ spl130_125), inference(avatar_split_clause, [], [f4207, f1721, f1406])).
fof(f4207, plain, (~ (e4 = op(e3, e0)) | ~ spl130_125), inference(forward_demodulation, [], [f596, f1723])).
fof(f4305, plain, (~ spl130_23 | ~ spl130_62 | ~ spl130_125 | spl130_295), inference(avatar_contradiction_clause, [], [f4304])).
fof(f4304, plain, ($false | (~ spl130_23 | ~ spl130_62 | ~ spl130_125 | spl130_295)), inference(subsumption_resolution, [], [f4299, f1459])).
fof(f4299, plain, (~ (e1 = op(e2, e2)) | (~ spl130_23 | ~ spl130_125 | spl130_295)), inference(backward_demodulation, [], [f4203, f1295])).
fof(f4203, plain, (~ (e1 = op(op(e4, e0), op(e4, e0))) | (~ spl130_125 | spl130_295)), inference(forward_demodulation, [], [f2836, f1723])).
fof(f4272, plain, (~ spl130_58 | ~ spl130_60), inference(avatar_contradiction_clause, [], [f4271])).
fof(f4271, plain, ($false | (~ spl130_58 | ~ spl130_60)), inference(subsumption_resolution, [], [f4270, f702])).
fof(f4270, plain, ((e2 = e4) | (~ spl130_58 | ~ spl130_60)), inference(backward_demodulation, [], [f1450, f1442])).
fof(f4268, plain, (~ spl130_66 | ~ spl130_67), inference(avatar_contradiction_clause, [], [f4267])).
fof(f4267, plain, ($false | (~ spl130_66 | ~ spl130_67)), inference(subsumption_resolution, [], [f4266, f694])).
fof(f4266, plain, ((e0 = e1) | (~ spl130_66 | ~ spl130_67)), inference(backward_demodulation, [], [f1480, f1476])).
fof(f4262, plain, (~ spl130_24 | ~ spl130_74), inference(avatar_split_clause, [], [f4259, f1507, f1297])).
fof(f4259, plain, (~ (e3 = op(e4, e0)) | ~ spl130_74), inference(backward_demodulation, [], [f602, f1509])).
fof(f4261, plain, (~ spl130_49 | ~ spl130_74), inference(avatar_split_clause, [], [f4258, f1507, f1402])).
fof(f4258, plain, (~ (e3 = op(e3, e0)) | ~ spl130_74), inference(backward_demodulation, [], [f601, f1509])).
fof(f4242, plain, (~ spl130_122 | ~ spl130_125), inference(avatar_contradiction_clause, [], [f4241])).
fof(f4241, plain, ($false | (~ spl130_122 | ~ spl130_125)), inference(subsumption_resolution, [], [f4240, f700])).
fof(f4240, plain, ((e1 = e4) | (~ spl130_122 | ~ spl130_125)), inference(backward_demodulation, [], [f1723, f1711])).
fof(f4235, plain, (~ spl130_28 | ~ spl130_129), inference(avatar_contradiction_clause, [], [f4234])).
fof(f4234, plain, ($false | (~ spl130_28 | ~ spl130_129)), inference(subsumption_resolution, [], [f4233, f702])).
fof(f4233, plain, ((e2 = e4) | (~ spl130_28 | ~ spl130_129)), inference(forward_demodulation, [], [f4216, f1316])).
fof(f4216, plain, ((e4 = op(e3, e4)) | ~ spl130_129), inference(backward_demodulation, [], [f541, f1740])).
fof(f4195, plain, (~ spl130_118 | ~ spl130_68), inference(avatar_split_clause, [], [f3932, f1482, f1692])).
fof(f3932, plain, (~ (e2 = op(e0, e1)) | ~ spl130_68), inference(forward_demodulation, [], [f605, f1484])).
fof(f4175, plain, (~ spl130_16 | ~ spl130_20), inference(avatar_contradiction_clause, [], [f4174])).
fof(f4174, plain, ($false | (~ spl130_16 | ~ spl130_20)), inference(subsumption_resolution, [], [f4171, f697])).
fof(f4171, plain, ((e0 = e4) | (~ spl130_16 | ~ spl130_20)), inference(backward_demodulation, [], [f1282, f1266])).
fof(f4155, plain, (~ spl130_56 | ~ spl130_60), inference(avatar_contradiction_clause, [], [f4154])).
fof(f4154, plain, ($false | (~ spl130_56 | ~ spl130_60)), inference(subsumption_resolution, [], [f4153, f697])).
fof(f4153, plain, ((e0 = e4) | (~ spl130_56 | ~ spl130_60)), inference(backward_demodulation, [], [f1450, f1434])).
fof(f4131, plain, (~ spl130_123 | ~ spl130_125), inference(avatar_contradiction_clause, [], [f4130])).
fof(f4130, plain, ($false | (~ spl130_123 | ~ spl130_125)), inference(subsumption_resolution, [], [f4129, f702])).
fof(f4129, plain, ((e2 = e4) | (~ spl130_123 | ~ spl130_125)), inference(forward_demodulation, [], [f1723, f1715])).
fof(f4125, plain, (~ spl130_54 | ~ spl130_128), inference(avatar_contradiction_clause, [], [f4124])).
fof(f4124, plain, ($false | (~ spl130_54 | ~ spl130_128)), inference(subsumption_resolution, [], [f4123, f703])).
fof(f4123, plain, ((e3 = e4) | (~ spl130_54 | ~ spl130_128)), inference(forward_demodulation, [], [f4101, f1425])).
fof(f4101, plain, ((e4 = op(e2, e4)) | ~ spl130_128), inference(backward_demodulation, [], [f541, f1736])).
fof(f4088, plain, (~ spl130_117 | ~ spl130_107), inference(avatar_split_clause, [], [f4087, f1646, f1688])).
fof(f4087, plain, (~ (e1 = op(e0, e1)) | ~ spl130_107), inference(forward_demodulation, [], [f649, f1648])).
fof(f4080, plain, (~ spl130_73 | ~ spl130_123), inference(avatar_split_clause, [], [f4079, f1713, f1503])).
fof(f4079, plain, (~ (e2 = op(e2, e0)) | ~ spl130_123), inference(forward_demodulation, [], [f595, f1715])).
fof(f4051, plain, (~ spl130_33 | ~ spl130_28), inference(avatar_split_clause, [], [f4050, f1314, f1335])).
fof(f4050, plain, (~ (e2 = op(e3, e3)) | ~ spl130_28), inference(forward_demodulation, [], [f683, f1316])).
fof(f4043, plain, (~ spl130_51 | ~ spl130_54), inference(avatar_contradiction_clause, [], [f4042])).
fof(f4042, plain, ($false | (~ spl130_51 | ~ spl130_54)), inference(subsumption_resolution, [], [f4041, f696])).
fof(f4041, plain, ((e0 = e3) | (~ spl130_51 | ~ spl130_54)), inference(forward_demodulation, [], [f1425, f1413])).
fof(f4040, plain, (~ spl130_57 | ~ spl130_60), inference(avatar_contradiction_clause, [], [f4039])).
fof(f4039, plain, ($false | (~ spl130_57 | ~ spl130_60)), inference(subsumption_resolution, [], [f4038, f700])).
fof(f4038, plain, ((e1 = e4) | (~ spl130_57 | ~ spl130_60)), inference(forward_demodulation, [], [f1450, f1438])).
fof(f4022, plain, (~ spl130_114 | ~ spl130_115), inference(avatar_contradiction_clause, [], [f4021])).
fof(f4021, plain, ($false | (~ spl130_114 | ~ spl130_115)), inference(subsumption_resolution, [], [f4020, f703])).
fof(f4020, plain, ((e3 = e4) | (~ spl130_114 | ~ spl130_115)), inference(forward_demodulation, [], [f1681, f1677])).
fof(f4012, plain, (~ spl130_2 | ~ spl130_75 | ~ spl130_123 | spl130_295), inference(avatar_contradiction_clause, [], [f4011])).
fof(f4011, plain, ($false | (~ spl130_2 | ~ spl130_75 | ~ spl130_123 | spl130_295)), inference(subsumption_resolution, [], [f4010, f1207])).
fof(f4010, plain, (~ (e1 = op(e4, e4)) | (~ spl130_75 | ~ spl130_123 | spl130_295)), inference(forward_demodulation, [], [f4003, f1513])).
fof(f4003, plain, (~ (e1 = op(op(e2, e0), op(e2, e0))) | (~ spl130_123 | spl130_295)), inference(backward_demodulation, [], [f2836, f1715])).
fof(f4008, plain, (~ spl130_108 | ~ spl130_123), inference(avatar_split_clause, [], [f3999, f1713, f1650])).
fof(f3999, plain, (~ (e2 = op(e0, e3)) | ~ spl130_123), inference(backward_demodulation, [], [f646, f1715])).
fof(f3994, plain, (~ spl130_122 | ~ spl130_31 | spl130_278), inference(avatar_split_clause, [], [f3844, f2745, f1327, f1709])).
fof(f3844, plain, (~ (op(e0, e0) = e1) | (~ spl130_31 | spl130_278)), inference(forward_demodulation, [], [f2747, f1329])).
fof(f3962, plain, (~ spl130_37 | ~ spl130_40), inference(avatar_contradiction_clause, [], [f3961])).
fof(f3961, plain, ($false | (~ spl130_37 | ~ spl130_40)), inference(subsumption_resolution, [], [f3960, f700])).
fof(f3960, plain, ((e1 = e4) | (~ spl130_37 | ~ spl130_40)), inference(forward_demodulation, [], [f1366, f1354])).
fof(f3947, plain, (~ spl130_30 | ~ spl130_80), inference(avatar_split_clause, [], [f3946, f1532, f1322])).
fof(f3946, plain, (~ (e4 = op(e3, e4)) | ~ spl130_80), inference(backward_demodulation, [], [f639, f1534])).
fof(f3942, plain, (~ spl130_108 | ~ spl130_110), inference(avatar_contradiction_clause, [], [f3941])).
fof(f3941, plain, ($false | (~ spl130_108 | ~ spl130_110)), inference(subsumption_resolution, [], [f3940, f702])).
fof(f3940, plain, ((e2 = e4) | (~ spl130_108 | ~ spl130_110)), inference(backward_demodulation, [], [f1660, f1652])).
fof(f3875, plain, (~ spl130_65 | ~ spl130_75), inference(avatar_split_clause, [], [f3874, f1511, f1469])).
fof(f3874, plain, (~ (e4 = op(e2, e2)) | ~ spl130_75), inference(forward_demodulation, [], [f665, f1513])).
fof(f3831, plain, (~ spl130_29 | ~ spl130_104), inference(avatar_split_clause, [], [f3830, f1633, f1318])).
fof(f3830, plain, (~ (e3 = op(e3, e4)) | ~ spl130_104), inference(forward_demodulation, [], [f636, f1635])).
fof(f3827, plain, (~ spl130_11 | ~ spl130_14), inference(avatar_contradiction_clause, [], [f3826])).
fof(f3826, plain, ($false | (~ spl130_11 | ~ spl130_14)), inference(subsumption_resolution, [], [f3824, f696])).
fof(f3824, plain, ((e0 = e3) | (~ spl130_11 | ~ spl130_14)), inference(backward_demodulation, [], [f1257, f1245])).
fof(f1245, plain, ((e0 = op(e4, e2)) | ~ spl130_11), inference(avatar_component_clause, [], [f1243])).
fof(f1257, plain, ((e3 = op(e4, e2)) | ~ spl130_14), inference(avatar_component_clause, [], [f1255])).
fof(f3807, plain, (~ spl130_48 | ~ spl130_49), inference(avatar_contradiction_clause, [], [f3806])).
fof(f3806, plain, ($false | (~ spl130_48 | ~ spl130_49)), inference(subsumption_resolution, [], [f3805, f701])).
fof(f3805, plain, ((e2 = e3) | (~ spl130_48 | ~ spl130_49)), inference(forward_demodulation, [], [f1404, f1400])).
fof(f3803, plain, (~ spl130_51 | ~ spl130_53), inference(avatar_contradiction_clause, [], [f3802])).
fof(f3802, plain, ($false | (~ spl130_51 | ~ spl130_53)), inference(subsumption_resolution, [], [f3801, f695])).
fof(f3801, plain, ((e0 = e2) | (~ spl130_51 | ~ spl130_53)), inference(backward_demodulation, [], [f1421, f1413])).
fof(f1421, plain, ((e2 = op(e2, e4)) | ~ spl130_53), inference(avatar_component_clause, [], [f1419])).
fof(f3791, plain, (~ spl130_84 | ~ spl130_85), inference(avatar_contradiction_clause, [], [f3790])).
fof(f3790, plain, ($false | (~ spl130_84 | ~ spl130_85)), inference(subsumption_resolution, [], [f3789, f703])).
fof(f3789, plain, ((e3 = e4) | (~ spl130_84 | ~ spl130_85)), inference(backward_demodulation, [], [f1555, f1551])).
fof(f3787, plain, (~ spl130_38 | ~ spl130_88), inference(avatar_split_clause, [], [f3785, f1566, f1356])).
fof(f3785, plain, (~ (e2 = op(e3, e2)) | ~ spl130_88), inference(backward_demodulation, [], [f619, f1568])).
fof(f3780, plain, (~ spl130_87 | ~ spl130_92), inference(avatar_split_clause, [], [f3772, f1583, f1562])).
fof(f3772, plain, (~ (e1 = op(e1, e2)) | ~ spl130_92), inference(backward_demodulation, [], [f658, f1585])).
fof(f3766, plain, (~ spl130_105 | ~ spl130_110), inference(avatar_split_clause, [], [f3765, f1658, f1637])).
fof(f3765, plain, (~ (e4 = op(e0, e4)) | ~ spl130_110), inference(backward_demodulation, [], [f653, f1660])).
fof(f3763, plain, (~ spl130_13 | ~ spl130_113), inference(avatar_split_clause, [], [f3756, f1671, f1251])).
fof(f3756, plain, (~ (e2 = op(e4, e2)) | ~ spl130_113), inference(backward_demodulation, [], [f617, f1673])).
fof(f3762, plain, (~ spl130_38 | ~ spl130_113), inference(avatar_split_clause, [], [f3755, f1671, f1356])).
fof(f3755, plain, (~ (e2 = op(e3, e2)) | ~ spl130_113), inference(backward_demodulation, [], [f616, f1673])).
fof(f3753, plain, (~ spl130_116 | ~ spl130_117), inference(avatar_contradiction_clause, [], [f3752])).
fof(f3752, plain, ($false | (~ spl130_116 | ~ spl130_117)), inference(subsumption_resolution, [], [f3751, f694])).
fof(f3751, plain, ((e0 = e1) | (~ spl130_116 | ~ spl130_117)), inference(backward_demodulation, [], [f1690, f1686])).
fof(f3696, plain, (~ spl130_89 | ~ spl130_79), inference(avatar_split_clause, [], [f3693, f1528, f1570])).
fof(f3693, plain, (~ (e3 = op(e1, e2)) | ~ spl130_79), inference(backward_demodulation, [], [f662, f1530])).
fof(f3690, plain, (~ spl130_77 | ~ spl130_97), inference(avatar_split_clause, [], [f3689, f1604, f1520])).
fof(f3689, plain, (~ (e1 = op(e1, e4)) | ~ spl130_97), inference(forward_demodulation, [], [f657, f1606])).
fof(f3674, plain, (~ spl130_14 | ~ spl130_65 | spl130_283), inference(avatar_split_clause, [], [f3673, f2770, f1469, f1255])).
fof(f3673, plain, (~ (e3 = op(e4, e2)) | (~ spl130_65 | spl130_283)), inference(forward_demodulation, [], [f2772, f1471])).
fof(f3657, plain, (~ spl130_38 | ~ spl130_48), inference(avatar_split_clause, [], [f3656, f1398, f1356])).
fof(f3656, plain, (~ (e2 = op(e3, e2)) | ~ spl130_48), inference(forward_demodulation, [], [f675, f1400])).
fof(f3655, plain, (~ spl130_47 | ~ spl130_32), inference(avatar_split_clause, [], [f3654, f1331, f1394])).
fof(f3654, plain, (~ (e1 = op(e3, e0)) | ~ spl130_32), inference(forward_demodulation, [], [f676, f1333])).
fof(f3651, plain, (~ spl130_73 | ~ spl130_48), inference(avatar_split_clause, [], [f3650, f1398, f1503])).
fof(f3650, plain, (~ (e2 = op(e2, e0)) | ~ spl130_48), inference(forward_demodulation, [], [f601, f1400])).
fof(f3625, plain, (~ spl130_83 | ~ spl130_8), inference(avatar_split_clause, [], [f3620, f1230, f1545])).
fof(f3620, plain, (~ (e2 = op(e1, e3)) | ~ spl130_8), inference(backward_demodulation, [], [f630, f1232])).
fof(f3584, plain, (~ spl130_48 | ~ spl130_50), inference(avatar_contradiction_clause, [], [f3583])).
fof(f3583, plain, ($false | (~ spl130_48 | ~ spl130_50)), inference(subsumption_resolution, [], [f3582, f702])).
fof(f3582, plain, ((e2 = e4) | (~ spl130_48 | ~ spl130_50)), inference(backward_demodulation, [], [f1408, f1400])).
fof(f3575, plain, (~ spl130_61 | ~ spl130_65), inference(avatar_contradiction_clause, [], [f3574])).
fof(f3574, plain, ($false | (~ spl130_61 | ~ spl130_65)), inference(subsumption_resolution, [], [f3573, f697])).
fof(f3573, plain, ((e0 = e4) | (~ spl130_61 | ~ spl130_65)), inference(forward_demodulation, [], [f1471, f1455])).
fof(f3565, plain, (~ spl130_57 | ~ spl130_72), inference(avatar_split_clause, [], [f3562, f1499, f1436])).
fof(f3562, plain, (~ (e1 = op(e2, e3)) | ~ spl130_72), inference(backward_demodulation, [], [f666, f1501])).
fof(f3564, plain, (~ spl130_67 | ~ spl130_72), inference(avatar_split_clause, [], [f3561, f1499, f1478])).
fof(f3561, plain, (~ (e1 = op(e2, e1)) | ~ spl130_72), inference(backward_demodulation, [], [f664, f1501])).
fof(f3557, plain, (~ spl130_83 | ~ spl130_85), inference(avatar_contradiction_clause, [], [f3556])).
fof(f3556, plain, ($false | (~ spl130_83 | ~ spl130_85)), inference(subsumption_resolution, [], [f3555, f702])).
fof(f3555, plain, ((e2 = e4) | (~ spl130_83 | ~ spl130_85)), inference(backward_demodulation, [], [f1555, f1547])).
fof(f3533, plain, (~ spl130_87 | ~ spl130_112), inference(avatar_split_clause, [], [f3529, f1667, f1562])).
fof(f3529, plain, (~ (e1 = op(e1, e2)) | ~ spl130_112), inference(backward_demodulation, [], [f614, f1669])).
fof(f3527, plain, (~ spl130_42 | ~ spl130_117), inference(avatar_split_clause, [], [f3523, f1688, f1373])).
fof(f3523, plain, (~ (e1 = op(e3, e1)) | ~ spl130_117), inference(backward_demodulation, [], [f606, f1690])).
fof(f3520, plain, (~ spl130_18 | ~ spl130_127), inference(avatar_contradiction_clause, [], [f3519])).
fof(f3519, plain, ($false | (~ spl130_18 | ~ spl130_127)), inference(subsumption_resolution, [], [f3518, f702])).
fof(f3518, plain, ((e2 = e4) | (~ spl130_18 | ~ spl130_127)), inference(forward_demodulation, [], [f3498, f1274])).
fof(f3498, plain, ((e4 = op(e4, e1)) | ~ spl130_127), inference(backward_demodulation, [], [f542, f1732])).
fof(f3517, plain, (~ spl130_77 | ~ spl130_127), inference(avatar_contradiction_clause, [], [f3516])).
fof(f3516, plain, ($false | (~ spl130_77 | ~ spl130_127)), inference(subsumption_resolution, [], [f3515, f700])).
fof(f3515, plain, ((e1 = e4) | (~ spl130_77 | ~ spl130_127)), inference(forward_demodulation, [], [f3497, f1522])).
fof(f3497, plain, ((e4 = op(e1, e4)) | ~ spl130_127), inference(backward_demodulation, [], [f541, f1732])).
fof(f3484, plain, (~ spl130_118 | ~ spl130_18), inference(avatar_split_clause, [], [f3483, f1272, f1692])).
fof(f3483, plain, (~ (e2 = op(e0, e1)) | ~ spl130_18), inference(forward_demodulation, [], [f607, f1274])).
fof(f3470, plain, (~ spl130_114 | ~ spl130_61 | spl130_283), inference(avatar_split_clause, [], [f3232, f2770, f1453, f1675])).
fof(f3232, plain, (~ (e3 = op(e0, e2)) | (~ spl130_61 | spl130_283)), inference(backward_demodulation, [], [f2772, f1455])).
fof(f3469, plain, (~ spl130_115 | ~ spl130_61 | spl130_287), inference(avatar_split_clause, [], [f3233, f2792, f1453, f1679])).
fof(f3233, plain, (~ (e4 = op(e0, e2)) | (~ spl130_61 | spl130_287)), inference(backward_demodulation, [], [f2794, f1455])).
fof(f2794, plain, (~ (e4 = op(op(e2, e2), e2)) | spl130_287), inference(avatar_component_clause, [], [f2792])).
fof(f3468, plain, (~ spl130_106 | ~ spl130_6), inference(avatar_split_clause, [], [f3467, f1222, f1642])).
fof(f3467, plain, (~ (e0 = op(e0, e3)) | ~ spl130_6), inference(forward_demodulation, [], [f627, f1224])).
fof(f3465, plain, (~ spl130_103 | ~ spl130_105), inference(avatar_contradiction_clause, [], [f3464])).
fof(f3464, plain, ($false | (~ spl130_103 | ~ spl130_105)), inference(subsumption_resolution, [], [f3463, f702])).
fof(f3463, plain, ((e2 = e4) | (~ spl130_103 | ~ spl130_105)), inference(forward_demodulation, [], [f1639, f1631])).
fof(f3429, plain, (~ spl130_66 | ~ spl130_61), inference(avatar_split_clause, [], [f3428, f1453, f1474])).
fof(f3428, plain, (~ (e0 = op(e2, e1)) | ~ spl130_61), inference(forward_demodulation, [], [f668, f1455])).
fof(f3410, plain, (~ spl130_31 | ~ spl130_26), inference(avatar_split_clause, [], [f3409, f1306, f1327])).
fof(f3409, plain, (~ (e0 = op(e3, e3)) | ~ spl130_26), inference(forward_demodulation, [], [f683, f1308])).
fof(f3408, plain, (~ spl130_26 | ~ spl130_29), inference(avatar_contradiction_clause, [], [f3407])).
fof(f3407, plain, ($false | (~ spl130_26 | ~ spl130_29)), inference(subsumption_resolution, [], [f3406, f696])).
fof(f3406, plain, ((e0 = e3) | (~ spl130_26 | ~ spl130_29)), inference(forward_demodulation, [], [f1320, f1308])).
fof(f1320, plain, ((e3 = op(e3, e4)) | ~ spl130_29), inference(avatar_component_clause, [], [f1318])).
fof(f3405, plain, (~ spl130_22 | ~ spl130_24), inference(avatar_contradiction_clause, [], [f3404])).
fof(f3404, plain, ($false | (~ spl130_22 | ~ spl130_24)), inference(subsumption_resolution, [], [f3403, f699])).
fof(f3403, plain, ((e1 = e3) | (~ spl130_22 | ~ spl130_24)), inference(forward_demodulation, [], [f1299, f1291])).
fof(f3394, plain, (~ spl130_6 | ~ spl130_8), inference(avatar_contradiction_clause, [], [f3393])).
fof(f3393, plain, ($false | (~ spl130_6 | ~ spl130_8)), inference(subsumption_resolution, [], [f3392, f695])).
fof(f3392, plain, ((e0 = e2) | (~ spl130_6 | ~ spl130_8)), inference(forward_demodulation, [], [f1232, f1224])).
fof(f3383, plain, (~ spl130_2 | ~ spl130_3), inference(avatar_contradiction_clause, [], [f3382])).
fof(f3382, plain, ($false | (~ spl130_2 | ~ spl130_3)), inference(subsumption_resolution, [], [f3381, f698])).
fof(f3381, plain, ((e1 = e2) | (~ spl130_2 | ~ spl130_3)), inference(backward_demodulation, [], [f1211, f1207])).
fof(f3375, plain, (~ spl130_4 | ~ spl130_26 | spl130_294), inference(avatar_contradiction_clause, [], [f3374])).
fof(f3374, plain, ($false | (~ spl130_4 | ~ spl130_26 | spl130_294)), inference(subsumption_resolution, [], [f3369, f1308])).
fof(f3369, plain, (~ (e0 = op(e3, e4)) | (~ spl130_4 | spl130_294)), inference(backward_demodulation, [], [f2830, f1215])).
fof(f3367, plain, (~ spl130_6 | ~ spl130_7), inference(avatar_contradiction_clause, [], [f3366])).
fof(f3366, plain, ($false | (~ spl130_6 | ~ spl130_7)), inference(subsumption_resolution, [], [f3365, f694])).
fof(f3365, plain, ((e0 = e1) | (~ spl130_6 | ~ spl130_7)), inference(backward_demodulation, [], [f1228, f1224])).
fof(f3361, plain, (~ spl130_8 | ~ spl130_9), inference(avatar_contradiction_clause, [], [f3360])).
fof(f3360, plain, ($false | (~ spl130_8 | ~ spl130_9)), inference(subsumption_resolution, [], [f3359, f701])).
fof(f3359, plain, ((e2 = e3) | (~ spl130_8 | ~ spl130_9)), inference(backward_demodulation, [], [f1236, f1232])).
fof(f1236, plain, ((e3 = op(e4, e3)) | ~ spl130_9), inference(avatar_component_clause, [], [f1234])).
fof(f3353, plain, (~ spl130_12 | ~ spl130_13), inference(avatar_contradiction_clause, [], [f3352])).
fof(f3352, plain, ($false | (~ spl130_12 | ~ spl130_13)), inference(subsumption_resolution, [], [f3351, f698])).
fof(f3351, plain, ((e1 = e2) | (~ spl130_12 | ~ spl130_13)), inference(backward_demodulation, [], [f1253, f1249])).
fof(f1253, plain, ((e2 = op(e4, e2)) | ~ spl130_13), inference(avatar_component_clause, [], [f1251])).
fof(f3350, plain, (~ spl130_13 | ~ spl130_14), inference(avatar_contradiction_clause, [], [f3349])).
fof(f3349, plain, ($false | (~ spl130_13 | ~ spl130_14)), inference(subsumption_resolution, [], [f3348, f701])).
fof(f3348, plain, ((e2 = e3) | (~ spl130_13 | ~ spl130_14)), inference(backward_demodulation, [], [f1257, f1253])).
fof(f3347, plain, (~ spl130_14 | ~ spl130_15), inference(avatar_contradiction_clause, [], [f3346])).
fof(f3346, plain, ($false | (~ spl130_14 | ~ spl130_15)), inference(subsumption_resolution, [], [f3345, f703])).
fof(f3345, plain, ((e3 = e4) | (~ spl130_14 | ~ spl130_15)), inference(backward_demodulation, [], [f1261, f1257])).
fof(f3344, plain, (~ spl130_5 | ~ spl130_15), inference(avatar_split_clause, [], [f3342, f1259, f1217])).
fof(f3342, plain, (~ (e4 = op(e4, e4)) | ~ spl130_15), inference(backward_demodulation, [], [f692, f1261])).
fof(f692, plain, ~ (op(e4, e2) = op(e4, e4)), inference(cnf_transformation, [], [f4])).
fof(f3337, plain, (~ spl130_18 | ~ spl130_19), inference(avatar_contradiction_clause, [], [f3336])).
fof(f3336, plain, ($false | (~ spl130_18 | ~ spl130_19)), inference(subsumption_resolution, [], [f3335, f701])).
fof(f3335, plain, ((e2 = e3) | (~ spl130_18 | ~ spl130_19)), inference(backward_demodulation, [], [f1278, f1274])).
fof(f3334, plain, (~ spl130_19 | ~ spl130_20), inference(avatar_contradiction_clause, [], [f3333])).
fof(f3333, plain, ($false | (~ spl130_19 | ~ spl130_20)), inference(subsumption_resolution, [], [f3332, f703])).
fof(f3332, plain, ((e3 = e4) | (~ spl130_19 | ~ spl130_20)), inference(backward_demodulation, [], [f1282, f1278])).
fof(f3330, plain, (~ spl130_10 | ~ spl130_20), inference(avatar_split_clause, [], [f3327, f1280, f1238])).
fof(f3327, plain, (~ (e4 = op(e4, e3)) | ~ spl130_20), inference(backward_demodulation, [], [f689, f1282])).
fof(f3325, plain, (~ spl130_22 | ~ spl130_23), inference(avatar_contradiction_clause, [], [f3324])).
fof(f3324, plain, ($false | (~ spl130_22 | ~ spl130_23)), inference(subsumption_resolution, [], [f3323, f698])).
fof(f3323, plain, ((e1 = e2) | (~ spl130_22 | ~ spl130_23)), inference(backward_demodulation, [], [f1295, f1291])).
fof(f3322, plain, (~ spl130_23 | ~ spl130_24), inference(avatar_contradiction_clause, [], [f3321])).
fof(f3321, plain, ($false | (~ spl130_23 | ~ spl130_24)), inference(subsumption_resolution, [], [f3320, f701])).
fof(f3320, plain, ((e2 = e3) | (~ spl130_23 | ~ spl130_24)), inference(backward_demodulation, [], [f1299, f1295])).
fof(f3317, plain, (~ spl130_14 | ~ spl130_24), inference(avatar_split_clause, [], [f3312, f1297, f1255])).
fof(f3312, plain, (~ (e3 = op(e4, e2)) | ~ spl130_24), inference(backward_demodulation, [], [f685, f1299])).
fof(f3316, plain, (~ spl130_19 | ~ spl130_24), inference(avatar_split_clause, [], [f3311, f1297, f1276])).
fof(f3311, plain, (~ (e3 = op(e4, e1)) | ~ spl130_24), inference(backward_demodulation, [], [f684, f1299])).
fof(f3307, plain, (~ spl130_27 | ~ spl130_28), inference(avatar_contradiction_clause, [], [f3306])).
fof(f3306, plain, ($false | (~ spl130_27 | ~ spl130_28)), inference(subsumption_resolution, [], [f3305, f698])).
fof(f3305, plain, ((e1 = e2) | (~ spl130_27 | ~ spl130_28)), inference(backward_demodulation, [], [f1316, f1312])).
fof(f3277, plain, (~ spl130_31 | ~ spl130_41), inference(avatar_split_clause, [], [f3272, f1369, f1327])).
fof(f3272, plain, (~ (e0 = op(e3, e3)) | ~ spl130_41), inference(backward_demodulation, [], [f679, f1371])).
fof(f3267, plain, (~ spl130_40 | ~ spl130_50), inference(avatar_split_clause, [], [f3262, f1406, f1364])).
fof(f3262, plain, (~ (e4 = op(e3, e2)) | ~ spl130_50), inference(backward_demodulation, [], [f675, f1408])).
fof(f3259, plain, (~ spl130_52 | ~ spl130_53), inference(avatar_contradiction_clause, [], [f3258])).
fof(f3258, plain, ($false | (~ spl130_52 | ~ spl130_53)), inference(subsumption_resolution, [], [f3257, f698])).
fof(f3257, plain, ((e1 = e2) | (~ spl130_52 | ~ spl130_53)), inference(backward_demodulation, [], [f1421, f1417])).
fof(f3246, plain, (~ spl130_51 | ~ spl130_56), inference(avatar_split_clause, [], [f3243, f1432, f1411])).
fof(f3243, plain, (~ (e0 = op(e2, e4)) | ~ spl130_56), inference(backward_demodulation, [], [f673, f1434])).
fof(f3210, plain, (~ spl130_56 | ~ spl130_71), inference(avatar_split_clause, [], [f3203, f1495, f1432])).
fof(f3203, plain, (~ (e0 = op(e2, e3)) | ~ spl130_71), inference(backward_demodulation, [], [f666, f1497])).
fof(f3206, plain, (~ spl130_46 | ~ spl130_71), inference(avatar_split_clause, [], [f3199, f1495, f1390])).
fof(f3199, plain, (~ (e0 = op(e3, e0)) | ~ spl130_71), inference(backward_demodulation, [], [f601, f1497])).
fof(f3195, plain, (~ spl130_78 | ~ spl130_79), inference(avatar_contradiction_clause, [], [f3194])).
fof(f3194, plain, ($false | (~ spl130_78 | ~ spl130_79)), inference(subsumption_resolution, [], [f3193, f701])).
fof(f3193, plain, ((e2 = e3) | (~ spl130_78 | ~ spl130_79)), inference(backward_demodulation, [], [f1530, f1526])).
fof(f3192, plain, (~ spl130_79 | ~ spl130_80), inference(avatar_contradiction_clause, [], [f3191])).
fof(f3191, plain, ($false | (~ spl130_79 | ~ spl130_80)), inference(subsumption_resolution, [], [f3190, f703])).
fof(f3190, plain, ((e3 = e4) | (~ spl130_79 | ~ spl130_80)), inference(backward_demodulation, [], [f1534, f1530])).
fof(f3170, plain, (~ spl130_36 | ~ spl130_86), inference(avatar_split_clause, [], [f3164, f1558, f1348])).
fof(f3164, plain, (~ (e0 = op(e3, e2)) | ~ spl130_86), inference(backward_demodulation, [], [f619, f1560])).
fof(f3155, plain, (~ spl130_41 | ~ spl130_91), inference(avatar_split_clause, [], [f3143, f1579, f1369])).
fof(f3143, plain, (~ (e0 = op(e3, e1)) | ~ spl130_91), inference(backward_demodulation, [], [f609, f1581])).
fof(f3136, plain, (~ spl130_46 | ~ spl130_96), inference(avatar_split_clause, [], [f3128, f1600, f1390])).
fof(f3128, plain, (~ (e0 = op(e3, e0)) | ~ spl130_96), inference(backward_demodulation, [], [f599, f1602])).
fof(f3135, plain, (~ spl130_71 | ~ spl130_96), inference(avatar_split_clause, [], [f3127, f1600, f1495])).
fof(f3127, plain, (~ (e0 = op(e2, e0)) | ~ spl130_96), inference(backward_demodulation, [], [f598, f1602])).
fof(f3126, plain, (~ spl130_102 | ~ spl130_103), inference(avatar_contradiction_clause, [], [f3125])).
fof(f3125, plain, ($false | (~ spl130_102 | ~ spl130_103)), inference(subsumption_resolution, [], [f3124, f698])).
fof(f3124, plain, ((e1 = e2) | (~ spl130_102 | ~ spl130_103)), inference(backward_demodulation, [], [f1631, f1627])).
fof(f3123, plain, (~ spl130_103 | ~ spl130_104), inference(avatar_contradiction_clause, [], [f3122])).
fof(f3122, plain, ($false | (~ spl130_103 | ~ spl130_104)), inference(subsumption_resolution, [], [f3121, f701])).
fof(f3121, plain, ((e2 = e3) | (~ spl130_103 | ~ spl130_104)), inference(backward_demodulation, [], [f1635, f1631])).
fof(f3120, plain, (~ spl130_104 | ~ spl130_105), inference(avatar_contradiction_clause, [], [f3119])).
fof(f3119, plain, ($false | (~ spl130_104 | ~ spl130_105)), inference(subsumption_resolution, [], [f3118, f703])).
fof(f3118, plain, ((e3 = e4) | (~ spl130_104 | ~ spl130_105)), inference(backward_demodulation, [], [f1639, f1635])).
fof(f3084, plain, (~ spl130_101 | ~ spl130_116), inference(avatar_split_clause, [], [f3077, f1684, f1621])).
fof(f3077, plain, (~ (e0 = op(e0, e4)) | ~ spl130_116), inference(backward_demodulation, [], [f650, f1686])).
fof(f3083, plain, (~ spl130_106 | ~ spl130_116), inference(avatar_split_clause, [], [f3076, f1684, f1642])).
fof(f3076, plain, (~ (e0 = op(e0, e3)) | ~ spl130_116), inference(backward_demodulation, [], [f649, f1686])).
fof(f3070, plain, (~ spl130_101 | ~ spl130_121), inference(avatar_split_clause, [], [f3058, f1705, f1621])).
fof(f3058, plain, (~ (e0 = op(e0, e4)) | ~ spl130_121), inference(backward_demodulation, [], [f647, f1707])).
fof(f3069, plain, (~ spl130_106 | ~ spl130_121), inference(avatar_split_clause, [], [f3057, f1705, f1642])).
fof(f3057, plain, (~ (e0 = op(e0, e3)) | ~ spl130_121), inference(backward_demodulation, [], [f646, f1707])).
fof(f3026, plain, (~ spl130_318 | ~ spl130_283 | ~ spl130_62 | ~ spl130_303), inference(avatar_split_clause, [], [f1199, f2875, f1457, f2770, f2994])).
fof(f1199, plain, (~ (e0 = op(op(e2, e2), op(e2, e2))) | ~ (e1 = op(e2, e2)) | ~ (e3 = op(op(e2, e2), e2)) | ~ (e4 = op(op(op(e2, e2), e2), op(op(e2, e2), e2)))), inference(cnf_transformation, [], [f246])).
fof(f246, plain, (~ (e0 = op(op(e2, e2), op(e2, e2))) | ~ (e1 = op(e2, e2)) | ~ (e3 = op(op(e2, e2), e2)) | ~ (e4 = op(op(op(e2, e2), e2), op(op(e2, e2), e2)))), inference(ennf_transformation, [], [f126])).
fof(f126, plain, ~ ((e0 = op(op(e2, e2), op(e2, e2))) & (e1 = op(e2, e2)) & (e3 = op(op(e2, e2), e2)) & (e4 = op(op(op(e2, e2), e2), op(op(e2, e2), e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG099+1.p', ax126)).
fof(f3021, plain, (~ spl130_320 | ~ spl130_301 | ~ spl130_122 | ~ spl130_302), inference(avatar_split_clause, [], [f1194, f2869, f1709, f2864, f3006])).
fof(f1194, plain, (~ (e2 = op(op(e0, e0), op(e0, e0))) | ~ (op(e0, e0) = e1) | ~ (e3 = op(op(e0, e0), e0)) | ~ (e4 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), inference(cnf_transformation, [], [f241])).
fof(f241, plain, (~ (e2 = op(op(e0, e0), op(e0, e0))) | ~ (op(e0, e0) = e1) | ~ (e3 = op(op(e0, e0), e0)) | ~ (e4 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), inference(ennf_transformation, [], [f121])).
fof(f121, plain, ~ ((e2 = op(op(e0, e0), op(e0, e0))) & (op(e0, e0) = e1) & (e3 = op(op(e0, e0), e0)) & (e4 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG099+1.p', ax121)).
fof(f3019, plain, (~ spl130_319 | ~ spl130_277 | ~ spl130_31 | ~ spl130_278), inference(avatar_split_clause, [], [f1192, f2745, f1327, f2740, f3000])).
fof(f1192, plain, (~ (e1 = op(op(e3, e3), op(e3, e3))) | ~ (e0 = op(e3, e3)) | ~ (e2 = op(op(e3, e3), e3)) | ~ (e4 = op(op(op(e3, e3), e3), op(op(e3, e3), e3)))), inference(cnf_transformation, [], [f239])).
fof(f239, plain, (~ (e1 = op(op(e3, e3), op(e3, e3))) | ~ (e0 = op(e3, e3)) | ~ (e2 = op(op(e3, e3), e3)) | ~ (e4 = op(op(op(e3, e3), e3), op(op(e3, e3), e3)))), inference(ennf_transformation, [], [f119])).
fof(f119, plain, ~ ((e1 = op(op(e3, e3), op(e3, e3))) & (e0 = op(e3, e3)) & (e2 = op(op(e3, e3), e3)) & (e4 = op(op(op(e3, e3), e3), op(op(e3, e3), e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG099+1.p', ax119)).
fof(f3016, plain, (~ spl130_320 | ~ spl130_296 | ~ spl130_124 | ~ spl130_311), inference(avatar_split_clause, [], [f1189, f2929, f1717, f2838, f3006])).
fof(f1189, plain, (~ (e1 = op(op(e0, e0), op(e0, e0))) | ~ (op(e0, e0) = e3) | ~ (e2 = op(op(e0, e0), e0)) | ~ (e4 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), inference(cnf_transformation, [], [f236])).
fof(f236, plain, (~ (e1 = op(op(e0, e0), op(e0, e0))) | ~ (op(e0, e0) = e3) | ~ (e2 = op(op(e0, e0), e0)) | ~ (e4 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), inference(ennf_transformation, [], [f116])).
fof(f116, plain, ~ ((e1 = op(op(e0, e0), op(e0, e0))) & (op(e0, e0) = e3) & (e2 = op(op(e0, e0), e0)) & (e4 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG099+1.p', ax116)).
fof(f3010, plain, (~ spl130_320 | ~ spl130_310 | ~ spl130_124 | ~ spl130_302), inference(avatar_split_clause, [], [f1183, f2869, f1717, f2918, f3006])).
fof(f1183, plain, (~ (e2 = op(op(e0, e0), op(e0, e0))) | ~ (op(e0, e0) = e3) | ~ (e1 = op(op(e0, e0), e0)) | ~ (e4 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), inference(cnf_transformation, [], [f230])).
fof(f230, plain, (~ (e2 = op(op(e0, e0), op(e0, e0))) | ~ (op(e0, e0) = e3) | ~ (e1 = op(op(e0, e0), e0)) | ~ (e4 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), inference(ennf_transformation, [], [f110])).
fof(f110, plain, ~ ((e2 = op(op(e0, e0), op(e0, e0))) & (op(e0, e0) = e3) & (e1 = op(op(e0, e0), e0)) & (e4 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG099+1.p', ax110)).
fof(f3009, plain, (~ spl130_320 | ~ spl130_310 | ~ spl130_123 | ~ spl130_298), inference(avatar_split_clause, [], [f1182, f2847, f1713, f2918, f3006])).
fof(f1182, plain, (~ (e3 = op(op(e0, e0), op(e0, e0))) | ~ (op(e0, e0) = e2) | ~ (e1 = op(op(e0, e0), e0)) | ~ (e4 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), inference(cnf_transformation, [], [f229])).
fof(f229, plain, (~ (e3 = op(op(e0, e0), op(e0, e0))) | ~ (op(e0, e0) = e2) | ~ (e1 = op(op(e0, e0), e0)) | ~ (e4 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), inference(ennf_transformation, [], [f109])).
fof(f109, plain, ~ ((e3 = op(op(e0, e0), op(e0, e0))) & (op(e0, e0) = e2) & (e1 = op(op(e0, e0), e0)) & (e4 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG099+1.p', ax109)).
fof(f3004, plain, (~ spl130_319 | ~ spl130_292 | ~ spl130_33 | ~ spl130_278), inference(avatar_split_clause, [], [f1181, f2745, f1335, f2818, f3000])).
fof(f1181, plain, (~ (e1 = op(op(e3, e3), op(e3, e3))) | ~ (e2 = op(e3, e3)) | ~ (e0 = op(op(e3, e3), e3)) | ~ (e4 = op(op(op(e3, e3), e3), op(op(e3, e3), e3)))), inference(cnf_transformation, [], [f228])).
fof(f228, plain, (~ (e1 = op(op(e3, e3), op(e3, e3))) | ~ (e2 = op(e3, e3)) | ~ (e0 = op(op(e3, e3), e3)) | ~ (e4 = op(op(op(e3, e3), e3), op(op(e3, e3), e3)))), inference(ennf_transformation, [], [f108])).
fof(f108, plain, ~ ((e1 = op(op(e3, e3), op(e3, e3))) & (e2 = op(e3, e3)) & (e0 = op(op(e3, e3), e3)) & (e4 = op(op(op(e3, e3), e3), op(op(e3, e3), e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG099+1.p', ax108)).
fof(f2998, plain, (~ spl130_318 | ~ spl130_290 | ~ spl130_64 | ~ spl130_284), inference(avatar_split_clause, [], [f1179, f2775, f1465, f2808, f2994])).
fof(f1179, plain, (~ (e1 = op(op(e2, e2), op(e2, e2))) | ~ (e3 = op(e2, e2)) | ~ (e0 = op(op(e2, e2), e2)) | ~ (e4 = op(op(op(e2, e2), e2), op(op(e2, e2), e2)))), inference(cnf_transformation, [], [f226])).
fof(f226, plain, (~ (e1 = op(op(e2, e2), op(e2, e2))) | ~ (e3 = op(e2, e2)) | ~ (e0 = op(op(e2, e2), e2)) | ~ (e4 = op(op(op(e2, e2), e2), op(op(e2, e2), e2)))), inference(ennf_transformation, [], [f106])).
fof(f106, plain, ~ ((e1 = op(op(e2, e2), op(e2, e2))) & (e3 = op(e2, e2)) & (e0 = op(op(e2, e2), e2)) & (e4 = op(op(op(e2, e2), e2), op(op(e2, e2), e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG099+1.p', ax106)).
fof(f2982, plain, (~ spl130_316 | ~ spl130_304 | ~ spl130_123 | ~ spl130_311), inference(avatar_split_clause, [], [f1171, f2929, f1713, f2882, f2966])).
fof(f1171, plain, (~ (e1 = op(op(e0, e0), op(e0, e0))) | ~ (op(e0, e0) = e2) | ~ (e4 = op(op(e0, e0), e0)) | ~ (e3 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), inference(cnf_transformation, [], [f218])).
fof(f218, plain, (~ (e1 = op(op(e0, e0), op(e0, e0))) | ~ (op(e0, e0) = e2) | ~ (e4 = op(op(e0, e0), e0)) | ~ (e3 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), inference(ennf_transformation, [], [f98])).
fof(f98, plain, ~ ((e1 = op(op(e0, e0), op(e0, e0))) & (op(e0, e0) = e2) & (e4 = op(op(e0, e0), e0)) & (e3 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG099+1.p', ax98)).
fof(f2981, plain, (~ spl130_316 | ~ spl130_304 | ~ spl130_122 | ~ spl130_302), inference(avatar_split_clause, [], [f1170, f2869, f1709, f2882, f2966])).
fof(f1170, plain, (~ (e2 = op(op(e0, e0), op(e0, e0))) | ~ (op(e0, e0) = e1) | ~ (e4 = op(op(e0, e0), e0)) | ~ (e3 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), inference(cnf_transformation, [], [f217])).
fof(f217, plain, (~ (e2 = op(op(e0, e0), op(e0, e0))) | ~ (op(e0, e0) = e1) | ~ (e4 = op(op(e0, e0), e0)) | ~ (e3 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), inference(ennf_transformation, [], [f97])).
fof(f97, plain, ~ ((e2 = op(op(e0, e0), op(e0, e0))) & (op(e0, e0) = e1) & (e4 = op(op(e0, e0), e0)) & (e3 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG099+1.p', ax97)).
fof(f2976, plain, (~ spl130_316 | ~ spl130_296 | ~ spl130_125 | ~ spl130_311), inference(avatar_split_clause, [], [f1165, f2929, f1721, f2838, f2966])).
fof(f1165, plain, (~ (e1 = op(op(e0, e0), op(e0, e0))) | ~ (op(e0, e0) = e4) | ~ (e2 = op(op(e0, e0), e0)) | ~ (e3 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), inference(cnf_transformation, [], [f212])).
fof(f212, plain, (~ (e1 = op(op(e0, e0), op(e0, e0))) | ~ (op(e0, e0) = e4) | ~ (e2 = op(op(e0, e0), e0)) | ~ (e3 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), inference(ennf_transformation, [], [f92])).
fof(f92, plain, ~ ((e1 = op(op(e0, e0), op(e0, e0))) & (op(e0, e0) = e4) & (e2 = op(op(e0, e0), e0)) & (e3 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG099+1.p', ax92)).
fof(f2975, plain, (~ spl130_316 | ~ spl130_296 | ~ spl130_122 | ~ spl130_297), inference(avatar_split_clause, [], [f1164, f2842, f1709, f2838, f2966])).
fof(f1164, plain, (~ (e4 = op(op(e0, e0), op(e0, e0))) | ~ (op(e0, e0) = e1) | ~ (e2 = op(op(e0, e0), e0)) | ~ (e3 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), inference(cnf_transformation, [], [f211])).
fof(f211, plain, (~ (e4 = op(op(e0, e0), op(e0, e0))) | ~ (op(e0, e0) = e1) | ~ (e2 = op(op(e0, e0), e0)) | ~ (e3 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), inference(ennf_transformation, [], [f91])).
fof(f91, plain, ~ ((e4 = op(op(e0, e0), op(e0, e0))) & (op(e0, e0) = e1) & (e2 = op(op(e0, e0), e0)) & (e3 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG099+1.p', ax91)).
fof(f2974, plain, (~ spl130_315 | ~ spl130_270 | ~ spl130_3 | ~ spl130_300), inference(avatar_split_clause, [], [f1163, f2859, f1209, f2708, f2960])).
fof(f1163, plain, (~ (e0 = op(op(e4, e4), op(e4, e4))) | ~ (e2 = op(e4, e4)) | ~ (e1 = op(op(e4, e4), e4)) | ~ (e3 = op(op(op(e4, e4), e4), op(op(e4, e4), e4)))), inference(cnf_transformation, [], [f210])).
fof(f210, plain, (~ (e0 = op(op(e4, e4), op(e4, e4))) | ~ (e2 = op(e4, e4)) | ~ (e1 = op(op(e4, e4), e4)) | ~ (e3 = op(op(op(e4, e4), e4), op(op(e4, e4), e4)))), inference(ennf_transformation, [], [f90])).
fof(f90, plain, ~ ((e0 = op(op(e4, e4), op(e4, e4))) & (e2 = op(e4, e4)) & (e1 = op(op(e4, e4), e4)) & (e3 = op(op(op(e4, e4), e4), op(op(e4, e4), e4)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG099+1.p', ax90)).
fof(f2969, plain, (~ spl130_316 | ~ spl130_310 | ~ spl130_123 | ~ spl130_297), inference(avatar_split_clause, [], [f1158, f2842, f1713, f2918, f2966])).
fof(f1158, plain, (~ (e4 = op(op(e0, e0), op(e0, e0))) | ~ (op(e0, e0) = e2) | ~ (e1 = op(op(e0, e0), e0)) | ~ (e3 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), inference(cnf_transformation, [], [f205])).
fof(f205, plain, (~ (e4 = op(op(e0, e0), op(e0, e0))) | ~ (op(e0, e0) = e2) | ~ (e1 = op(op(e0, e0), e0)) | ~ (e3 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), inference(ennf_transformation, [], [f85])).
fof(f85, plain, ~ ((e4 = op(op(e0, e0), op(e0, e0))) & (op(e0, e0) = e2) & (e1 = op(op(e0, e0), e0)) & (e3 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG099+1.p', ax85)).
fof(f2964, plain, (~ spl130_315 | ~ spl130_294 | ~ spl130_3 | ~ spl130_280), inference(avatar_split_clause, [], [f1157, f2755, f1209, f2828, f2960])).
fof(f1157, plain, (~ (e1 = op(op(e4, e4), op(e4, e4))) | ~ (e2 = op(e4, e4)) | ~ (e0 = op(op(e4, e4), e4)) | ~ (e3 = op(op(op(e4, e4), e4), op(op(e4, e4), e4)))), inference(cnf_transformation, [], [f204])).
fof(f204, plain, (~ (e1 = op(op(e4, e4), op(e4, e4))) | ~ (e2 = op(e4, e4)) | ~ (e0 = op(op(e4, e4), e4)) | ~ (e3 = op(op(op(e4, e4), e4), op(op(e4, e4), e4)))), inference(ennf_transformation, [], [f84])).
fof(f84, plain, ~ ((e1 = op(op(e4, e4), op(e4, e4))) & (e2 = op(e4, e4)) & (e0 = op(op(e4, e4), e4)) & (e3 = op(op(op(e4, e4), e4), op(op(e4, e4), e4)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG099+1.p', ax84)).
fof(f2958, plain, (~ spl130_314 | ~ spl130_290 | ~ spl130_65 | ~ spl130_284), inference(avatar_split_clause, [], [f1155, f2775, f1469, f2808, f2954])).
fof(f1155, plain, (~ (e1 = op(op(e2, e2), op(e2, e2))) | ~ (e4 = op(e2, e2)) | ~ (e0 = op(op(e2, e2), e2)) | ~ (e3 = op(op(op(e2, e2), e2), op(op(e2, e2), e2)))), inference(cnf_transformation, [], [f202])).
fof(f202, plain, (~ (e1 = op(op(e2, e2), op(e2, e2))) | ~ (e4 = op(e2, e2)) | ~ (e0 = op(op(e2, e2), e2)) | ~ (e3 = op(op(op(e2, e2), e2), op(op(e2, e2), e2)))), inference(ennf_transformation, [], [f82])).
fof(f82, plain, ~ ((e1 = op(op(e2, e2), op(e2, e2))) & (e4 = op(e2, e2)) & (e0 = op(op(e2, e2), e2)) & (e3 = op(op(op(e2, e2), e2), op(op(e2, e2), e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG099+1.p', ax82)).
fof(f2945, plain, (~ spl130_307 | ~ spl130_288 | ~ spl130_31 | ~ spl130_278), inference(avatar_split_clause, [], [f1150, f2745, f1327, f2798, f2902])).
fof(f1150, plain, (~ (e1 = op(op(e3, e3), op(e3, e3))) | ~ (e0 = op(e3, e3)) | ~ (e4 = op(op(e3, e3), e3)) | ~ (e2 = op(op(op(e3, e3), e3), op(op(e3, e3), e3)))), inference(cnf_transformation, [], [f197])).
fof(f197, plain, (~ (e1 = op(op(e3, e3), op(e3, e3))) | ~ (e0 = op(e3, e3)) | ~ (e4 = op(op(e3, e3), e3)) | ~ (e2 = op(op(op(e3, e3), e3), op(op(e3, e3), e3)))), inference(ennf_transformation, [], [f77])).
fof(f77, plain, ~ ((e1 = op(op(e3, e3), op(e3, e3))) & (e0 = op(e3, e3)) & (e4 = op(op(e3, e3), e3)) & (e2 = op(op(op(e3, e3), e3), op(op(e3, e3), e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG099+1.p', ax77)).
fof(f2942, plain, (~ spl130_309 | ~ spl130_304 | ~ spl130_124 | ~ spl130_311), inference(avatar_split_clause, [], [f1147, f2929, f1717, f2882, f2914])).
fof(f1147, plain, (~ (e1 = op(op(e0, e0), op(e0, e0))) | ~ (op(e0, e0) = e3) | ~ (e4 = op(op(e0, e0), e0)) | ~ (e2 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), inference(cnf_transformation, [], [f194])).
fof(f194, plain, (~ (e1 = op(op(e0, e0), op(e0, e0))) | ~ (op(e0, e0) = e3) | ~ (e4 = op(op(e0, e0), e0)) | ~ (e2 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), inference(ennf_transformation, [], [f74])).
fof(f74, plain, ~ ((e1 = op(op(e0, e0), op(e0, e0))) & (op(e0, e0) = e3) & (e4 = op(op(e0, e0), e0)) & (e2 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG099+1.p', ax74)).
fof(f2932, plain, (~ spl130_309 | ~ spl130_301 | ~ spl130_125 | ~ spl130_311), inference(avatar_split_clause, [], [f1141, f2929, f1721, f2864, f2914])).
fof(f1141, plain, (~ (e1 = op(op(e0, e0), op(e0, e0))) | ~ (op(e0, e0) = e4) | ~ (e3 = op(op(e0, e0), e0)) | ~ (e2 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), inference(cnf_transformation, [], [f188])).
fof(f188, plain, (~ (e1 = op(op(e0, e0), op(e0, e0))) | ~ (op(e0, e0) = e4) | ~ (e3 = op(op(e0, e0), e0)) | ~ (e2 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), inference(ennf_transformation, [], [f68])).
fof(f68, plain, ~ ((e1 = op(op(e0, e0), op(e0, e0))) & (op(e0, e0) = e4) & (e3 = op(op(e0, e0), e0)) & (e2 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG099+1.p', ax68)).
fof(f2927, plain, (~ spl130_309 | ~ spl130_301 | ~ spl130_122 | ~ spl130_297), inference(avatar_split_clause, [], [f1140, f2842, f1709, f2864, f2914])).
fof(f1140, plain, (~ (e4 = op(op(e0, e0), op(e0, e0))) | ~ (op(e0, e0) = e1) | ~ (e3 = op(op(e0, e0), e0)) | ~ (e2 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), inference(cnf_transformation, [], [f187])).
fof(f187, plain, (~ (e4 = op(op(e0, e0), op(e0, e0))) | ~ (op(e0, e0) = e1) | ~ (e3 = op(op(e0, e0), e0)) | ~ (e2 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), inference(ennf_transformation, [], [f67])).
fof(f67, plain, ~ ((e4 = op(op(e0, e0), op(e0, e0))) & (op(e0, e0) = e1) & (e3 = op(op(e0, e0), e0)) & (e2 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG099+1.p', ax67)).
fof(f2922, plain, (~ spl130_309 | ~ spl130_310 | ~ spl130_125 | ~ spl130_298), inference(avatar_split_clause, [], [f1135, f2847, f1721, f2918, f2914])).
fof(f1135, plain, (~ (e3 = op(op(e0, e0), op(e0, e0))) | ~ (op(e0, e0) = e4) | ~ (e1 = op(op(e0, e0), e0)) | ~ (e2 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), inference(cnf_transformation, [], [f182])).
fof(f182, plain, (~ (e3 = op(op(e0, e0), op(e0, e0))) | ~ (op(e0, e0) = e4) | ~ (e1 = op(op(e0, e0), e0)) | ~ (e2 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), inference(ennf_transformation, [], [f62])).
fof(f62, plain, ~ ((e3 = op(op(e0, e0), op(e0, e0))) & (op(e0, e0) = e4) & (e1 = op(op(e0, e0), e0)) & (e2 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG099+1.p', ax62)).
fof(f2921, plain, (~ spl130_309 | ~ spl130_310 | ~ spl130_124 | ~ spl130_297), inference(avatar_split_clause, [], [f1134, f2842, f1717, f2918, f2914])).
fof(f1134, plain, (~ (e4 = op(op(e0, e0), op(e0, e0))) | ~ (op(e0, e0) = e3) | ~ (e1 = op(op(e0, e0), e0)) | ~ (e2 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), inference(cnf_transformation, [], [f181])).
fof(f181, plain, (~ (e4 = op(op(e0, e0), op(e0, e0))) | ~ (op(e0, e0) = e3) | ~ (e1 = op(op(e0, e0), e0)) | ~ (e2 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), inference(ennf_transformation, [], [f61])).
fof(f61, plain, ~ ((e4 = op(op(e0, e0), op(e0, e0))) & (op(e0, e0) = e3) & (e1 = op(op(e0, e0), e0)) & (e2 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG099+1.p', ax61)).
fof(f2911, plain, (~ spl130_308 | ~ spl130_294 | ~ spl130_2 | ~ spl130_271), inference(avatar_split_clause, [], [f1132, f2712, f1205, f2828, f2908])).
fof(f1132, plain, (~ (e3 = op(op(e4, e4), op(e4, e4))) | ~ (e1 = op(e4, e4)) | ~ (e0 = op(op(e4, e4), e4)) | ~ (e2 = op(op(op(e4, e4), e4), op(op(e4, e4), e4)))), inference(cnf_transformation, [], [f179])).
fof(f179, plain, (~ (e3 = op(op(e4, e4), op(e4, e4))) | ~ (e1 = op(e4, e4)) | ~ (e0 = op(op(e4, e4), e4)) | ~ (e2 = op(op(op(e4, e4), e4), op(op(e4, e4), e4)))), inference(ennf_transformation, [], [f59])).
fof(f59, plain, ~ ((e3 = op(op(e4, e4), op(e4, e4))) & (e1 = op(e4, e4)) & (e0 = op(op(e4, e4), e4)) & (e2 = op(op(op(e4, e4), e4), op(op(e4, e4), e4)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG099+1.p', ax59)).
fof(f2900, plain, (~ spl130_305 | ~ spl130_306 | ~ spl130_95 | ~ spl130_276), inference(avatar_split_clause, [], [f1129, f2735, f1595, f2896, f2892])).
fof(f1129, plain, (~ (e3 = op(op(e1, e1), op(e1, e1))) | ~ (e4 = op(e1, e1)) | ~ (e0 = op(op(e1, e1), e1)) | ~ (e2 = op(op(op(e1, e1), e1), op(op(e1, e1), e1)))), inference(cnf_transformation, [], [f176])).
fof(f176, plain, (~ (e3 = op(op(e1, e1), op(e1, e1))) | ~ (e4 = op(e1, e1)) | ~ (e0 = op(op(e1, e1), e1)) | ~ (e2 = op(op(op(e1, e1), e1), op(op(e1, e1), e1)))), inference(ennf_transformation, [], [f56])).
fof(f56, plain, ~ ((e3 = op(op(e1, e1), op(e1, e1))) & (e4 = op(e1, e1)) & (e0 = op(op(e1, e1), e1)) & (e2 = op(op(op(e1, e1), e1), op(op(e1, e1), e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG099+1.p', ax56)).
fof(f2887, plain, (~ spl130_289 | ~ spl130_287 | ~ spl130_61 | ~ spl130_264), inference(avatar_split_clause, [], [f1124, f2681, f1453, f2792, f2804])).
fof(f1124, plain, (~ (e3 = op(op(e2, e2), op(e2, e2))) | ~ (e0 = op(e2, e2)) | ~ (e4 = op(op(e2, e2), e2)) | ~ (e1 = op(op(op(e2, e2), e2), op(op(e2, e2), e2)))), inference(cnf_transformation, [], [f171])).
fof(f171, plain, (~ (e3 = op(op(e2, e2), op(e2, e2))) | ~ (e0 = op(e2, e2)) | ~ (e4 = op(op(e2, e2), e2)) | ~ (e1 = op(op(op(e2, e2), e2), op(op(e2, e2), e2)))), inference(ennf_transformation, [], [f51])).
fof(f51, plain, ~ ((e3 = op(op(e2, e2), op(e2, e2))) & (e0 = op(e2, e2)) & (e4 = op(op(e2, e2), e2)) & (e1 = op(op(op(e2, e2), e2), op(op(e2, e2), e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG099+1.p', ax51)).
fof(f2885, plain, (~ spl130_295 | ~ spl130_304 | ~ spl130_123 | ~ spl130_298), inference(avatar_split_clause, [], [f1122, f2847, f1713, f2882, f2834])).
fof(f1122, plain, (~ (e3 = op(op(e0, e0), op(e0, e0))) | ~ (op(e0, e0) = e2) | ~ (e4 = op(op(e0, e0), e0)) | ~ (e1 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), inference(cnf_transformation, [], [f169])).
fof(f169, plain, (~ (e3 = op(op(e0, e0), op(e0, e0))) | ~ (op(e0, e0) = e2) | ~ (e4 = op(op(e0, e0), e0)) | ~ (e1 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), inference(ennf_transformation, [], [f49])).
fof(f49, plain, ~ ((e3 = op(op(e0, e0), op(e0, e0))) & (op(e0, e0) = e2) & (e4 = op(op(e0, e0), e0)) & (e1 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG099+1.p', ax49)).
fof(f2879, plain, (~ spl130_293 | ~ spl130_285 | ~ spl130_1 | ~ spl130_272), inference(avatar_split_clause, [], [f1120, f2717, f1201, f2780, f2824])).
fof(f1120, plain, (~ (e2 = op(op(e4, e4), op(e4, e4))) | ~ (e0 = op(e4, e4)) | ~ (e3 = op(op(e4, e4), e4)) | ~ (e1 = op(op(op(e4, e4), e4), op(op(e4, e4), e4)))), inference(cnf_transformation, [], [f167])).
fof(f167, plain, (~ (e2 = op(op(e4, e4), op(e4, e4))) | ~ (e0 = op(e4, e4)) | ~ (e3 = op(op(e4, e4), e4)) | ~ (e1 = op(op(op(e4, e4), e4), op(op(e4, e4), e4)))), inference(ennf_transformation, [], [f47])).
fof(f47, plain, ~ ((e2 = op(op(e4, e4), op(e4, e4))) & (e0 = op(e4, e4)) & (e3 = op(op(e4, e4), e4)) & (e1 = op(op(op(e4, e4), e4), op(op(e4, e4), e4)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG099+1.p', ax47)).
fof(f2873, plain, (~ spl130_289 | ~ spl130_283 | ~ spl130_61 | ~ spl130_263), inference(avatar_split_clause, [], [f1118, f2676, f1453, f2770, f2804])).
fof(f1118, plain, (~ (e4 = op(op(e2, e2), op(e2, e2))) | ~ (e0 = op(e2, e2)) | ~ (e3 = op(op(e2, e2), e2)) | ~ (e1 = op(op(op(e2, e2), e2), op(op(e2, e2), e2)))), inference(cnf_transformation, [], [f165])).
fof(f165, plain, (~ (e4 = op(op(e2, e2), op(e2, e2))) | ~ (e0 = op(e2, e2)) | ~ (e3 = op(op(e2, e2), e2)) | ~ (e1 = op(op(op(e2, e2), e2), op(op(e2, e2), e2)))), inference(ennf_transformation, [], [f45])).
fof(f45, plain, ~ ((e4 = op(op(e2, e2), op(e2, e2))) & (e0 = op(e2, e2)) & (e3 = op(op(e2, e2), e2)) & (e1 = op(op(op(e2, e2), e2), op(op(e2, e2), e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG099+1.p', ax45)).
fof(f2850, plain, (~ spl130_295 | ~ spl130_296 | ~ spl130_125 | ~ spl130_298), inference(avatar_split_clause, [], [f1111, f2847, f1721, f2838, f2834])).
fof(f1111, plain, (~ (e3 = op(op(e0, e0), op(e0, e0))) | ~ (op(e0, e0) = e4) | ~ (e2 = op(op(e0, e0), e0)) | ~ (e1 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), inference(cnf_transformation, [], [f158])).
fof(f158, plain, (~ (e3 = op(op(e0, e0), op(e0, e0))) | ~ (op(e0, e0) = e4) | ~ (e2 = op(op(e0, e0), e0)) | ~ (e1 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), inference(ennf_transformation, [], [f38])).
fof(f38, plain, ~ ((e3 = op(op(e0, e0), op(e0, e0))) & (op(e0, e0) = e4) & (e2 = op(op(e0, e0), e0)) & (e1 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG099+1.p', ax38)).
fof(f2845, plain, (~ spl130_295 | ~ spl130_296 | ~ spl130_124 | ~ spl130_297), inference(avatar_split_clause, [], [f1110, f2842, f1717, f2838, f2834])).
fof(f1110, plain, (~ (e4 = op(op(e0, e0), op(e0, e0))) | ~ (op(e0, e0) = e3) | ~ (e2 = op(op(e0, e0), e0)) | ~ (e1 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), inference(cnf_transformation, [], [f157])).
fof(f157, plain, (~ (e4 = op(op(e0, e0), op(e0, e0))) | ~ (op(e0, e0) = e3) | ~ (e2 = op(op(e0, e0), e0)) | ~ (e1 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), inference(ennf_transformation, [], [f37])).
fof(f37, plain, ~ ((e4 = op(op(e0, e0), op(e0, e0))) & (op(e0, e0) = e3) & (e2 = op(op(e0, e0), e0)) & (e1 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG099+1.p', ax37)).
fof(f2832, plain, (~ spl130_293 | ~ spl130_294 | ~ spl130_4 | ~ spl130_272), inference(avatar_split_clause, [], [f1109, f2717, f1213, f2828, f2824])).
fof(f1109, plain, (~ (e2 = op(op(e4, e4), op(e4, e4))) | ~ (e3 = op(e4, e4)) | ~ (e0 = op(op(e4, e4), e4)) | ~ (e1 = op(op(op(e4, e4), e4), op(op(e4, e4), e4)))), inference(cnf_transformation, [], [f156])).
fof(f156, plain, (~ (e2 = op(op(e4, e4), op(e4, e4))) | ~ (e3 = op(e4, e4)) | ~ (e0 = op(op(e4, e4), e4)) | ~ (e1 = op(op(op(e4, e4), e4), op(op(e4, e4), e4)))), inference(ennf_transformation, [], [f36])).
fof(f36, plain, ~ ((e2 = op(op(e4, e4), op(e4, e4))) & (e3 = op(e4, e4)) & (e0 = op(op(e4, e4), e4)) & (e1 = op(op(op(e4, e4), e4), op(op(e4, e4), e4)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG099+1.p', ax36)).
fof(f2831, plain, (~ spl130_293 | ~ spl130_294 | ~ spl130_3 | ~ spl130_271), inference(avatar_split_clause, [], [f1108, f2712, f1209, f2828, f2824])).
fof(f1108, plain, (~ (e3 = op(op(e4, e4), op(e4, e4))) | ~ (e2 = op(e4, e4)) | ~ (e0 = op(op(e4, e4), e4)) | ~ (e1 = op(op(op(e4, e4), e4), op(op(e4, e4), e4)))), inference(cnf_transformation, [], [f155])).
fof(f155, plain, (~ (e3 = op(op(e4, e4), op(e4, e4))) | ~ (e2 = op(e4, e4)) | ~ (e0 = op(op(e4, e4), e4)) | ~ (e1 = op(op(op(e4, e4), e4), op(op(e4, e4), e4)))), inference(ennf_transformation, [], [f35])).
fof(f35, plain, ~ ((e3 = op(op(e4, e4), op(e4, e4))) & (e2 = op(e4, e4)) & (e0 = op(op(e4, e4), e4)) & (e1 = op(op(op(e4, e4), e4), op(op(e4, e4), e4)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG099+1.p', ax35)).
fof(f2790, plain, (~ spl130_273 | ~ spl130_286 | ~ spl130_94 | ~ spl130_282), inference(avatar_split_clause, [], [f1099, f2765, f1591, f2786, f2722])).
fof(f1099, plain, (~ (e2 = op(op(e1, e1), op(e1, e1))) | ~ (e3 = op(e1, e1)) | ~ (e4 = op(op(e1, e1), e1)) | ~ (e0 = op(op(op(e1, e1), e1), op(op(e1, e1), e1)))), inference(cnf_transformation, [], [f146])).
fof(f146, plain, (~ (e2 = op(op(e1, e1), op(e1, e1))) | ~ (e3 = op(e1, e1)) | ~ (e4 = op(op(e1, e1), e1)) | ~ (e0 = op(op(op(e1, e1), e1), op(op(e1, e1), e1)))), inference(ennf_transformation, [], [f26])).
fof(f26, plain, ~ ((e2 = op(op(e1, e1), op(e1, e1))) & (e3 = op(e1, e1)) & (e4 = op(op(e1, e1), e1)) & (e0 = op(op(op(e1, e1), e1), op(op(e1, e1), e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG099+1.p', ax26)).
fof(f2789, plain, (~ spl130_273 | ~ spl130_286 | ~ spl130_93 | ~ spl130_276), inference(avatar_split_clause, [], [f1098, f2735, f1587, f2786, f2722])).
fof(f1098, plain, (~ (e3 = op(op(e1, e1), op(e1, e1))) | ~ (e2 = op(e1, e1)) | ~ (e4 = op(op(e1, e1), e1)) | ~ (e0 = op(op(op(e1, e1), e1), op(op(e1, e1), e1)))), inference(cnf_transformation, [], [f145])).
fof(f145, plain, (~ (e3 = op(op(e1, e1), op(e1, e1))) | ~ (e2 = op(e1, e1)) | ~ (e4 = op(op(e1, e1), e1)) | ~ (e0 = op(op(op(e1, e1), e1), op(op(e1, e1), e1)))), inference(ennf_transformation, [], [f25])).
fof(f25, plain, ~ ((e3 = op(op(e1, e1), op(e1, e1))) & (e2 = op(e1, e1)) & (e4 = op(op(e1, e1), e1)) & (e0 = op(op(op(e1, e1), e1), op(op(e1, e1), e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG099+1.p', ax25)).
fof(f2778, plain, (~ spl130_261 | ~ spl130_283 | ~ spl130_65 | ~ spl130_284), inference(avatar_split_clause, [], [f1095, f2775, f1469, f2770, f2668])).
fof(f1095, plain, (~ (e1 = op(op(e2, e2), op(e2, e2))) | ~ (e4 = op(e2, e2)) | ~ (e3 = op(op(e2, e2), e2)) | ~ (e0 = op(op(op(e2, e2), e2), op(op(e2, e2), e2)))), inference(cnf_transformation, [], [f142])).
fof(f142, plain, (~ (e1 = op(op(e2, e2), op(e2, e2))) | ~ (e4 = op(e2, e2)) | ~ (e3 = op(op(e2, e2), e2)) | ~ (e0 = op(op(op(e2, e2), e2), op(op(e2, e2), e2)))), inference(ennf_transformation, [], [f22])).
fof(f22, plain, ~ ((e1 = op(op(e2, e2), op(e2, e2))) & (e4 = op(e2, e2)) & (e3 = op(op(e2, e2), e2)) & (e0 = op(op(op(e2, e2), e2), op(op(e2, e2), e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG099+1.p', ax22)).
fof(f2768, plain, (~ spl130_273 | ~ spl130_281 | ~ spl130_95 | ~ spl130_282), inference(avatar_split_clause, [], [f1093, f2765, f1595, f2760, f2722])).
fof(f1093, plain, (~ (e2 = op(op(e1, e1), op(e1, e1))) | ~ (e4 = op(e1, e1)) | ~ (e3 = op(op(e1, e1), e1)) | ~ (e0 = op(op(op(e1, e1), e1), op(op(e1, e1), e1)))), inference(cnf_transformation, [], [f140])).
fof(f140, plain, (~ (e2 = op(op(e1, e1), op(e1, e1))) | ~ (e4 = op(e1, e1)) | ~ (e3 = op(op(e1, e1), e1)) | ~ (e0 = op(op(op(e1, e1), e1), op(op(e1, e1), e1)))), inference(ennf_transformation, [], [f20])).
fof(f20, plain, ~ ((e2 = op(op(e1, e1), op(e1, e1))) & (e4 = op(e1, e1)) & (e3 = op(op(e1, e1), e1)) & (e0 = op(op(op(e1, e1), e1), op(op(e1, e1), e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG099+1.p', ax20)).
fof(f2738, plain, (~ spl130_273 | ~ spl130_274 | ~ spl130_95 | ~ spl130_276), inference(avatar_split_clause, [], [f1087, f2735, f1595, f2726, f2722])).
fof(f1087, plain, (~ (e3 = op(op(e1, e1), op(e1, e1))) | ~ (e4 = op(e1, e1)) | ~ (e2 = op(op(e1, e1), e1)) | ~ (e0 = op(op(op(e1, e1), e1), op(op(e1, e1), e1)))), inference(cnf_transformation, [], [f134])).
fof(f134, plain, (~ (e3 = op(op(e1, e1), op(e1, e1))) | ~ (e4 = op(e1, e1)) | ~ (e2 = op(op(e1, e1), e1)) | ~ (e0 = op(op(op(e1, e1), e1), op(op(e1, e1), e1)))), inference(ennf_transformation, [], [f14])).
fof(f14, plain, ~ ((e3 = op(op(e1, e1), op(e1, e1))) & (e4 = op(e1, e1)) & (e2 = op(op(e1, e1), e1)) & (e0 = op(op(op(e1, e1), e1), op(op(e1, e1), e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG099+1.p', ax14)).
fof(f2733, plain, (~ spl130_273 | ~ spl130_274 | ~ spl130_94 | ~ spl130_275), inference(avatar_split_clause, [], [f1086, f2730, f1591, f2726, f2722])).
fof(f1086, plain, (~ (e4 = op(op(e1, e1), op(e1, e1))) | ~ (e3 = op(e1, e1)) | ~ (e2 = op(op(e1, e1), e1)) | ~ (e0 = op(op(op(e1, e1), e1), op(op(e1, e1), e1)))), inference(cnf_transformation, [], [f133])).
fof(f133, plain, (~ (e4 = op(op(e1, e1), op(e1, e1))) | ~ (e3 = op(e1, e1)) | ~ (e2 = op(op(e1, e1), e1)) | ~ (e0 = op(op(op(e1, e1), e1), op(op(e1, e1), e1)))), inference(ennf_transformation, [], [f13])).
fof(f13, plain, ~ ((e4 = op(op(e1, e1), op(e1, e1))) & (e3 = op(e1, e1)) & (e2 = op(op(e1, e1), e1)) & (e0 = op(op(op(e1, e1), e1), op(op(e1, e1), e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG099+1.p', ax13)).
fof(f2702, plain, (~ spl130_265 | ~ spl130_266 | ~ spl130_35 | ~ spl130_268), inference(avatar_split_clause, [], [f1083, f2699, f1343, f2690, f2686])).
fof(f1083, plain, (~ (e2 = op(op(e3, e3), op(e3, e3))) | ~ (e4 = op(e3, e3)) | ~ (e1 = op(op(e3, e3), e3)) | ~ (e0 = op(op(op(e3, e3), e3), op(op(e3, e3), e3)))), inference(cnf_transformation, [], [f130])).
fof(f130, plain, (~ (e2 = op(op(e3, e3), op(e3, e3))) | ~ (e4 = op(e3, e3)) | ~ (e1 = op(op(e3, e3), e3)) | ~ (e0 = op(op(op(e3, e3), e3), op(op(e3, e3), e3)))), inference(ennf_transformation, [], [f10])).
fof(f10, plain, ~ ((e2 = op(op(e3, e3), op(e3, e3))) & (e4 = op(e3, e3)) & (e1 = op(op(e3, e3), e3)) & (e0 = op(op(op(e3, e3), e3), op(op(e3, e3), e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG099+1.p', ax10)).
fof(f2684, plain, (~ spl130_261 | ~ spl130_262 | ~ spl130_65 | ~ spl130_264), inference(avatar_split_clause, [], [f1081, f2681, f1469, f2672, f2668])).
fof(f1081, plain, (~ (e3 = op(op(e2, e2), op(e2, e2))) | ~ (e4 = op(e2, e2)) | ~ (e1 = op(op(e2, e2), e2)) | ~ (e0 = op(op(op(e2, e2), e2), op(op(e2, e2), e2)))), inference(cnf_transformation, [], [f128])).
fof(f128, plain, (~ (e3 = op(op(e2, e2), op(e2, e2))) | ~ (e4 = op(e2, e2)) | ~ (e1 = op(op(e2, e2), e2)) | ~ (e0 = op(op(op(e2, e2), e2), op(op(e2, e2), e2)))), inference(ennf_transformation, [], [f8])).
fof(f8, plain, ~ ((e3 = op(op(e2, e2), op(e2, e2))) & (e4 = op(e2, e2)) & (e1 = op(op(e2, e2), e2)) & (e0 = op(op(op(e2, e2), e2), op(op(e2, e2), e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG099+1.p', ax8)).
fof(f2666, plain, (spl130_235 | spl130_209 | spl130_183 | spl130_157 | spl130_131), inference(avatar_split_clause, [], [f1079, f1797, f1926, f2055, f2184, f2313])).
fof(f2313, plain, (spl130_235 <=> sP125), introduced(avatar_definition, [new_symbols(naming, [spl130_235])])).
fof(f2184, plain, (spl130_209 <=> sP126), introduced(avatar_definition, [new_symbols(naming, [spl130_209])])).
fof(f2055, plain, (spl130_183 <=> sP127), introduced(avatar_definition, [new_symbols(naming, [spl130_183])])).
fof(f1926, plain, (spl130_157 <=> sP128), introduced(avatar_definition, [new_symbols(naming, [spl130_157])])).
fof(f1797, plain, (spl130_131 <=> sP129), introduced(avatar_definition, [new_symbols(naming, [spl130_131])])).
fof(f1079, plain, (sP129 | sP128 | sP127 | sP126 | sP125), inference(cnf_transformation, [], [f377])).
fof(f377, plain, (sP129 | sP128 | sP127 | sP126 | sP125), inference(definition_folding, [], [f6, e376, e375, e374, e373, e372, e371, e370, e369, e368, e367, e366, e365, e364, e363, e362, e361, e360, e359, e358, e357, e356, e355, e354, e353, e352, e351, e350, e349, e348, e347, e346, e345, e344, e343, e342, e341, e340, e339, e338, e337, e336, e335, e334, e333, e332, e331, e330, e329, e328, e327, e326, e325, e324, e323, e322, e321, e320, e319, e318, e317, e316, e315, e314, e313, e312, e311, e310, e309, e308, e307, e306, e305, e304, e303, e302, e301, e300, e299, e298, e297, e296, e295, e294, e293, e292, e291, e290, e289, e288, e287, e286, e285, e284, e283, e282, e281, e280, e279, e278, e277, e276, e275, e274, e273, e272, e271, e270, e269, e268, e267, e266, e265, e264, e263, e262, e261, e260, e259, e258, e257, e256, e255, e254, e253, e252, e251, e250, e249, e248, e247])).
fof(f247, plain, ((~ (e0 = unit) & (e0 = op(e0, e0))) | ~ (e0 = op(e0, e0)) | ~ sP0), inference(usedef, [], [e247])).
fof(e247, plain, (sP0 <=> ((~ (e0 = unit) & (e0 = op(e0, e0))) | ~ (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f248, plain, ((~ (e0 = unit) & (op(e0, e0) = e1)) | ~ (e0 = op(e0, e1)) | ~ sP1), inference(usedef, [], [e248])).
fof(e248, plain, (sP1 <=> ((~ (e0 = unit) & (op(e0, e0) = e1)) | ~ (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f249, plain, ((~ (e0 = unit) & (op(e0, e0) = e2)) | ~ (e0 = op(e0, e2)) | ~ sP2), inference(usedef, [], [e249])).
fof(e249, plain, (sP2 <=> ((~ (e0 = unit) & (op(e0, e0) = e2)) | ~ (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f250, plain, ((~ (e0 = unit) & (op(e0, e0) = e3)) | ~ (e0 = op(e0, e3)) | ~ sP3), inference(usedef, [], [e250])).
fof(e250, plain, (sP3 <=> ((~ (e0 = unit) & (op(e0, e0) = e3)) | ~ (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f251, plain, ((~ (e0 = unit) & (op(e0, e0) = e4)) | ~ (e0 = op(e0, e4)) | ~ sP4), inference(usedef, [], [e251])).
fof(e251, plain, (sP4 <=> ((~ (e0 = unit) & (op(e0, e0) = e4)) | ~ (e0 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f252, plain, ((~ (e0 = unit) & (e0 = op(e1, e0))) | ~ (e0 = op(e1, e0)) | ~ sP5), inference(usedef, [], [e252])).
fof(e252, plain, (sP5 <=> ((~ (e0 = unit) & (e0 = op(e1, e0))) | ~ (e0 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f253, plain, ((~ (e0 = unit) & (e1 = op(e1, e0))) | ~ (e0 = op(e1, e1)) | ~ sP6), inference(usedef, [], [e253])).
fof(e253, plain, (sP6 <=> ((~ (e0 = unit) & (e1 = op(e1, e0))) | ~ (e0 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f254, plain, ((~ (e0 = unit) & (e2 = op(e1, e0))) | ~ (e0 = op(e1, e2)) | ~ sP7), inference(usedef, [], [e254])).
fof(e254, plain, (sP7 <=> ((~ (e0 = unit) & (e2 = op(e1, e0))) | ~ (e0 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f255, plain, ((~ (e0 = unit) & (e3 = op(e1, e0))) | ~ (e0 = op(e1, e3)) | ~ sP8), inference(usedef, [], [e255])).
fof(e255, plain, (sP8 <=> ((~ (e0 = unit) & (e3 = op(e1, e0))) | ~ (e0 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f256, plain, ((~ (e0 = unit) & (e4 = op(e1, e0))) | ~ (e0 = op(e1, e4)) | ~ sP9), inference(usedef, [], [e256])).
fof(e256, plain, (sP9 <=> ((~ (e0 = unit) & (e4 = op(e1, e0))) | ~ (e0 = op(e1, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f257, plain, ((~ (e0 = unit) & (e0 = op(e2, e0))) | ~ (e0 = op(e2, e0)) | ~ sP10), inference(usedef, [], [e257])).
fof(e257, plain, (sP10 <=> ((~ (e0 = unit) & (e0 = op(e2, e0))) | ~ (e0 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f258, plain, ((~ (e0 = unit) & (e1 = op(e2, e0))) | ~ (e0 = op(e2, e1)) | ~ sP11), inference(usedef, [], [e258])).
fof(e258, plain, (sP11 <=> ((~ (e0 = unit) & (e1 = op(e2, e0))) | ~ (e0 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f259, plain, ((~ (e0 = unit) & (e2 = op(e2, e0))) | ~ (e0 = op(e2, e2)) | ~ sP12), inference(usedef, [], [e259])).
fof(e259, plain, (sP12 <=> ((~ (e0 = unit) & (e2 = op(e2, e0))) | ~ (e0 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f260, plain, ((~ (e0 = unit) & (e3 = op(e2, e0))) | ~ (e0 = op(e2, e3)) | ~ sP13), inference(usedef, [], [e260])).
fof(e260, plain, (sP13 <=> ((~ (e0 = unit) & (e3 = op(e2, e0))) | ~ (e0 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f261, plain, ((~ (e0 = unit) & (e4 = op(e2, e0))) | ~ (e0 = op(e2, e4)) | ~ sP14), inference(usedef, [], [e261])).
fof(e261, plain, (sP14 <=> ((~ (e0 = unit) & (e4 = op(e2, e0))) | ~ (e0 = op(e2, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f262, plain, ((~ (e0 = unit) & (e0 = op(e3, e0))) | ~ (e0 = op(e3, e0)) | ~ sP15), inference(usedef, [], [e262])).
fof(e262, plain, (sP15 <=> ((~ (e0 = unit) & (e0 = op(e3, e0))) | ~ (e0 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f263, plain, ((~ (e0 = unit) & (e1 = op(e3, e0))) | ~ (e0 = op(e3, e1)) | ~ sP16), inference(usedef, [], [e263])).
fof(e263, plain, (sP16 <=> ((~ (e0 = unit) & (e1 = op(e3, e0))) | ~ (e0 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f264, plain, ((~ (e0 = unit) & (e2 = op(e3, e0))) | ~ (e0 = op(e3, e2)) | ~ sP17), inference(usedef, [], [e264])).
fof(e264, plain, (sP17 <=> ((~ (e0 = unit) & (e2 = op(e3, e0))) | ~ (e0 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f265, plain, ((~ (e0 = unit) & (e3 = op(e3, e0))) | ~ (e0 = op(e3, e3)) | ~ sP18), inference(usedef, [], [e265])).
fof(e265, plain, (sP18 <=> ((~ (e0 = unit) & (e3 = op(e3, e0))) | ~ (e0 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f266, plain, ((~ (e0 = unit) & (e4 = op(e3, e0))) | ~ (e0 = op(e3, e4)) | ~ sP19), inference(usedef, [], [e266])).
fof(e266, plain, (sP19 <=> ((~ (e0 = unit) & (e4 = op(e3, e0))) | ~ (e0 = op(e3, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP19])])).
fof(f267, plain, ((~ (e0 = unit) & (e0 = op(e4, e0))) | ~ (e0 = op(e4, e0)) | ~ sP20), inference(usedef, [], [e267])).
fof(e267, plain, (sP20 <=> ((~ (e0 = unit) & (e0 = op(e4, e0))) | ~ (e0 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP20])])).
fof(f268, plain, ((~ (e0 = unit) & (e1 = op(e4, e0))) | ~ (e0 = op(e4, e1)) | ~ sP21), inference(usedef, [], [e268])).
fof(e268, plain, (sP21 <=> ((~ (e0 = unit) & (e1 = op(e4, e0))) | ~ (e0 = op(e4, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP21])])).
fof(f269, plain, ((~ (e0 = unit) & (e2 = op(e4, e0))) | ~ (e0 = op(e4, e2)) | ~ sP22), inference(usedef, [], [e269])).
fof(e269, plain, (sP22 <=> ((~ (e0 = unit) & (e2 = op(e4, e0))) | ~ (e0 = op(e4, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP22])])).
fof(f270, plain, ((~ (e0 = unit) & (e3 = op(e4, e0))) | ~ (e0 = op(e4, e3)) | ~ sP23), inference(usedef, [], [e270])).
fof(e270, plain, (sP23 <=> ((~ (e0 = unit) & (e3 = op(e4, e0))) | ~ (e0 = op(e4, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP23])])).
fof(f271, plain, ((~ (e0 = unit) & (e4 = op(e4, e0))) | ~ (e0 = op(e4, e4)) | ~ sP24), inference(usedef, [], [e271])).
fof(e271, plain, (sP24 <=> ((~ (e0 = unit) & (e4 = op(e4, e0))) | ~ (e0 = op(e4, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP24])])).
fof(f272, plain, ((~ (e1 = unit) & (e0 = op(e0, e1))) | ~ (op(e0, e0) = e1) | ~ sP25), inference(usedef, [], [e272])).
fof(e272, plain, (sP25 <=> ((~ (e1 = unit) & (e0 = op(e0, e1))) | ~ (op(e0, e0) = e1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP25])])).
fof(f273, plain, ((~ (e1 = unit) & (e1 = op(e0, e1))) | ~ (e1 = op(e0, e1)) | ~ sP26), inference(usedef, [], [e273])).
fof(e273, plain, (sP26 <=> ((~ (e1 = unit) & (e1 = op(e0, e1))) | ~ (e1 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP26])])).
fof(f274, plain, ((~ (e1 = unit) & (e2 = op(e0, e1))) | ~ (e1 = op(e0, e2)) | ~ sP27), inference(usedef, [], [e274])).
fof(e274, plain, (sP27 <=> ((~ (e1 = unit) & (e2 = op(e0, e1))) | ~ (e1 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP27])])).
fof(f275, plain, ((~ (e1 = unit) & (e3 = op(e0, e1))) | ~ (e1 = op(e0, e3)) | ~ sP28), inference(usedef, [], [e275])).
fof(e275, plain, (sP28 <=> ((~ (e1 = unit) & (e3 = op(e0, e1))) | ~ (e1 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP28])])).
fof(f276, plain, ((~ (e1 = unit) & (e4 = op(e0, e1))) | ~ (e1 = op(e0, e4)) | ~ sP29), inference(usedef, [], [e276])).
fof(e276, plain, (sP29 <=> ((~ (e1 = unit) & (e4 = op(e0, e1))) | ~ (e1 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP29])])).
fof(f277, plain, ((~ (e1 = unit) & (e0 = op(e1, e1))) | ~ (e1 = op(e1, e0)) | ~ sP30), inference(usedef, [], [e277])).
fof(e277, plain, (sP30 <=> ((~ (e1 = unit) & (e0 = op(e1, e1))) | ~ (e1 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP30])])).
fof(f278, plain, ((~ (e1 = unit) & (e1 = op(e1, e1))) | ~ (e1 = op(e1, e1)) | ~ sP31), inference(usedef, [], [e278])).
fof(e278, plain, (sP31 <=> ((~ (e1 = unit) & (e1 = op(e1, e1))) | ~ (e1 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP31])])).
fof(f279, plain, ((~ (e1 = unit) & (e2 = op(e1, e1))) | ~ (e1 = op(e1, e2)) | ~ sP32), inference(usedef, [], [e279])).
fof(e279, plain, (sP32 <=> ((~ (e1 = unit) & (e2 = op(e1, e1))) | ~ (e1 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP32])])).
fof(f280, plain, ((~ (e1 = unit) & (e3 = op(e1, e1))) | ~ (e1 = op(e1, e3)) | ~ sP33), inference(usedef, [], [e280])).
fof(e280, plain, (sP33 <=> ((~ (e1 = unit) & (e3 = op(e1, e1))) | ~ (e1 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP33])])).
fof(f281, plain, ((~ (e1 = unit) & (e4 = op(e1, e1))) | ~ (e1 = op(e1, e4)) | ~ sP34), inference(usedef, [], [e281])).
fof(e281, plain, (sP34 <=> ((~ (e1 = unit) & (e4 = op(e1, e1))) | ~ (e1 = op(e1, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP34])])).
fof(f282, plain, ((~ (e1 = unit) & (e0 = op(e2, e1))) | ~ (e1 = op(e2, e0)) | ~ sP35), inference(usedef, [], [e282])).
fof(e282, plain, (sP35 <=> ((~ (e1 = unit) & (e0 = op(e2, e1))) | ~ (e1 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP35])])).
fof(f283, plain, ((~ (e1 = unit) & (e1 = op(e2, e1))) | ~ (e1 = op(e2, e1)) | ~ sP36), inference(usedef, [], [e283])).
fof(e283, plain, (sP36 <=> ((~ (e1 = unit) & (e1 = op(e2, e1))) | ~ (e1 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP36])])).
fof(f284, plain, ((~ (e1 = unit) & (e2 = op(e2, e1))) | ~ (e1 = op(e2, e2)) | ~ sP37), inference(usedef, [], [e284])).
fof(e284, plain, (sP37 <=> ((~ (e1 = unit) & (e2 = op(e2, e1))) | ~ (e1 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP37])])).
fof(f285, plain, ((~ (e1 = unit) & (e3 = op(e2, e1))) | ~ (e1 = op(e2, e3)) | ~ sP38), inference(usedef, [], [e285])).
fof(e285, plain, (sP38 <=> ((~ (e1 = unit) & (e3 = op(e2, e1))) | ~ (e1 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP38])])).
fof(f286, plain, ((~ (e1 = unit) & (e4 = op(e2, e1))) | ~ (e1 = op(e2, e4)) | ~ sP39), inference(usedef, [], [e286])).
fof(e286, plain, (sP39 <=> ((~ (e1 = unit) & (e4 = op(e2, e1))) | ~ (e1 = op(e2, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP39])])).
fof(f287, plain, ((~ (e1 = unit) & (e0 = op(e3, e1))) | ~ (e1 = op(e3, e0)) | ~ sP40), inference(usedef, [], [e287])).
fof(e287, plain, (sP40 <=> ((~ (e1 = unit) & (e0 = op(e3, e1))) | ~ (e1 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP40])])).
fof(f288, plain, ((~ (e1 = unit) & (e1 = op(e3, e1))) | ~ (e1 = op(e3, e1)) | ~ sP41), inference(usedef, [], [e288])).
fof(e288, plain, (sP41 <=> ((~ (e1 = unit) & (e1 = op(e3, e1))) | ~ (e1 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP41])])).
fof(f289, plain, ((~ (e1 = unit) & (e2 = op(e3, e1))) | ~ (e1 = op(e3, e2)) | ~ sP42), inference(usedef, [], [e289])).
fof(e289, plain, (sP42 <=> ((~ (e1 = unit) & (e2 = op(e3, e1))) | ~ (e1 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP42])])).
fof(f290, plain, ((~ (e1 = unit) & (e3 = op(e3, e1))) | ~ (e1 = op(e3, e3)) | ~ sP43), inference(usedef, [], [e290])).
fof(e290, plain, (sP43 <=> ((~ (e1 = unit) & (e3 = op(e3, e1))) | ~ (e1 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP43])])).
fof(f291, plain, ((~ (e1 = unit) & (e4 = op(e3, e1))) | ~ (e1 = op(e3, e4)) | ~ sP44), inference(usedef, [], [e291])).
fof(e291, plain, (sP44 <=> ((~ (e1 = unit) & (e4 = op(e3, e1))) | ~ (e1 = op(e3, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP44])])).
fof(f292, plain, ((~ (e1 = unit) & (e0 = op(e4, e1))) | ~ (e1 = op(e4, e0)) | ~ sP45), inference(usedef, [], [e292])).
fof(e292, plain, (sP45 <=> ((~ (e1 = unit) & (e0 = op(e4, e1))) | ~ (e1 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP45])])).
fof(f293, plain, ((~ (e1 = unit) & (e1 = op(e4, e1))) | ~ (e1 = op(e4, e1)) | ~ sP46), inference(usedef, [], [e293])).
fof(e293, plain, (sP46 <=> ((~ (e1 = unit) & (e1 = op(e4, e1))) | ~ (e1 = op(e4, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP46])])).
fof(f294, plain, ((~ (e1 = unit) & (e2 = op(e4, e1))) | ~ (e1 = op(e4, e2)) | ~ sP47), inference(usedef, [], [e294])).
fof(e294, plain, (sP47 <=> ((~ (e1 = unit) & (e2 = op(e4, e1))) | ~ (e1 = op(e4, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP47])])).
fof(f295, plain, ((~ (e1 = unit) & (e3 = op(e4, e1))) | ~ (e1 = op(e4, e3)) | ~ sP48), inference(usedef, [], [e295])).
fof(e295, plain, (sP48 <=> ((~ (e1 = unit) & (e3 = op(e4, e1))) | ~ (e1 = op(e4, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP48])])).
fof(f296, plain, ((~ (e1 = unit) & (e4 = op(e4, e1))) | ~ (e1 = op(e4, e4)) | ~ sP49), inference(usedef, [], [e296])).
fof(e296, plain, (sP49 <=> ((~ (e1 = unit) & (e4 = op(e4, e1))) | ~ (e1 = op(e4, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP49])])).
fof(f297, plain, ((~ (e2 = unit) & (e0 = op(e0, e2))) | ~ (op(e0, e0) = e2) | ~ sP50), inference(usedef, [], [e297])).
fof(e297, plain, (sP50 <=> ((~ (e2 = unit) & (e0 = op(e0, e2))) | ~ (op(e0, e0) = e2))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP50])])).
fof(f298, plain, ((~ (e2 = unit) & (e1 = op(e0, e2))) | ~ (e2 = op(e0, e1)) | ~ sP51), inference(usedef, [], [e298])).
fof(e298, plain, (sP51 <=> ((~ (e2 = unit) & (e1 = op(e0, e2))) | ~ (e2 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP51])])).
fof(f299, plain, ((~ (e2 = unit) & (e2 = op(e0, e2))) | ~ (e2 = op(e0, e2)) | ~ sP52), inference(usedef, [], [e299])).
fof(e299, plain, (sP52 <=> ((~ (e2 = unit) & (e2 = op(e0, e2))) | ~ (e2 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP52])])).
fof(f300, plain, ((~ (e2 = unit) & (e3 = op(e0, e2))) | ~ (e2 = op(e0, e3)) | ~ sP53), inference(usedef, [], [e300])).
fof(e300, plain, (sP53 <=> ((~ (e2 = unit) & (e3 = op(e0, e2))) | ~ (e2 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP53])])).
fof(f301, plain, ((~ (e2 = unit) & (e4 = op(e0, e2))) | ~ (e2 = op(e0, e4)) | ~ sP54), inference(usedef, [], [e301])).
fof(e301, plain, (sP54 <=> ((~ (e2 = unit) & (e4 = op(e0, e2))) | ~ (e2 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP54])])).
fof(f302, plain, ((~ (e2 = unit) & (e0 = op(e1, e2))) | ~ (e2 = op(e1, e0)) | ~ sP55), inference(usedef, [], [e302])).
fof(e302, plain, (sP55 <=> ((~ (e2 = unit) & (e0 = op(e1, e2))) | ~ (e2 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP55])])).
fof(f303, plain, ((~ (e2 = unit) & (e1 = op(e1, e2))) | ~ (e2 = op(e1, e1)) | ~ sP56), inference(usedef, [], [e303])).
fof(e303, plain, (sP56 <=> ((~ (e2 = unit) & (e1 = op(e1, e2))) | ~ (e2 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP56])])).
fof(f304, plain, ((~ (e2 = unit) & (e2 = op(e1, e2))) | ~ (e2 = op(e1, e2)) | ~ sP57), inference(usedef, [], [e304])).
fof(e304, plain, (sP57 <=> ((~ (e2 = unit) & (e2 = op(e1, e2))) | ~ (e2 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP57])])).
fof(f305, plain, ((~ (e2 = unit) & (e3 = op(e1, e2))) | ~ (e2 = op(e1, e3)) | ~ sP58), inference(usedef, [], [e305])).
fof(e305, plain, (sP58 <=> ((~ (e2 = unit) & (e3 = op(e1, e2))) | ~ (e2 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP58])])).
fof(f306, plain, ((~ (e2 = unit) & (e4 = op(e1, e2))) | ~ (e2 = op(e1, e4)) | ~ sP59), inference(usedef, [], [e306])).
fof(e306, plain, (sP59 <=> ((~ (e2 = unit) & (e4 = op(e1, e2))) | ~ (e2 = op(e1, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP59])])).
fof(f307, plain, ((~ (e2 = unit) & (e0 = op(e2, e2))) | ~ (e2 = op(e2, e0)) | ~ sP60), inference(usedef, [], [e307])).
fof(e307, plain, (sP60 <=> ((~ (e2 = unit) & (e0 = op(e2, e2))) | ~ (e2 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP60])])).
fof(f308, plain, ((~ (e2 = unit) & (e1 = op(e2, e2))) | ~ (e2 = op(e2, e1)) | ~ sP61), inference(usedef, [], [e308])).
fof(e308, plain, (sP61 <=> ((~ (e2 = unit) & (e1 = op(e2, e2))) | ~ (e2 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP61])])).
fof(f309, plain, ((~ (e2 = unit) & (e2 = op(e2, e2))) | ~ (e2 = op(e2, e2)) | ~ sP62), inference(usedef, [], [e309])).
fof(e309, plain, (sP62 <=> ((~ (e2 = unit) & (e2 = op(e2, e2))) | ~ (e2 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP62])])).
fof(f310, plain, ((~ (e2 = unit) & (e3 = op(e2, e2))) | ~ (e2 = op(e2, e3)) | ~ sP63), inference(usedef, [], [e310])).
fof(e310, plain, (sP63 <=> ((~ (e2 = unit) & (e3 = op(e2, e2))) | ~ (e2 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP63])])).
fof(f311, plain, ((~ (e2 = unit) & (e4 = op(e2, e2))) | ~ (e2 = op(e2, e4)) | ~ sP64), inference(usedef, [], [e311])).
fof(e311, plain, (sP64 <=> ((~ (e2 = unit) & (e4 = op(e2, e2))) | ~ (e2 = op(e2, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP64])])).
fof(f312, plain, ((~ (e2 = unit) & (e0 = op(e3, e2))) | ~ (e2 = op(e3, e0)) | ~ sP65), inference(usedef, [], [e312])).
fof(e312, plain, (sP65 <=> ((~ (e2 = unit) & (e0 = op(e3, e2))) | ~ (e2 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP65])])).
fof(f313, plain, ((~ (e2 = unit) & (e1 = op(e3, e2))) | ~ (e2 = op(e3, e1)) | ~ sP66), inference(usedef, [], [e313])).
fof(e313, plain, (sP66 <=> ((~ (e2 = unit) & (e1 = op(e3, e2))) | ~ (e2 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP66])])).
fof(f314, plain, ((~ (e2 = unit) & (e2 = op(e3, e2))) | ~ (e2 = op(e3, e2)) | ~ sP67), inference(usedef, [], [e314])).
fof(e314, plain, (sP67 <=> ((~ (e2 = unit) & (e2 = op(e3, e2))) | ~ (e2 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP67])])).
fof(f315, plain, ((~ (e2 = unit) & (e3 = op(e3, e2))) | ~ (e2 = op(e3, e3)) | ~ sP68), inference(usedef, [], [e315])).
fof(e315, plain, (sP68 <=> ((~ (e2 = unit) & (e3 = op(e3, e2))) | ~ (e2 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP68])])).
fof(f316, plain, ((~ (e2 = unit) & (e4 = op(e3, e2))) | ~ (e2 = op(e3, e4)) | ~ sP69), inference(usedef, [], [e316])).
fof(e316, plain, (sP69 <=> ((~ (e2 = unit) & (e4 = op(e3, e2))) | ~ (e2 = op(e3, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP69])])).
fof(f317, plain, ((~ (e2 = unit) & (e0 = op(e4, e2))) | ~ (e2 = op(e4, e0)) | ~ sP70), inference(usedef, [], [e317])).
fof(e317, plain, (sP70 <=> ((~ (e2 = unit) & (e0 = op(e4, e2))) | ~ (e2 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP70])])).
fof(f318, plain, ((~ (e2 = unit) & (e1 = op(e4, e2))) | ~ (e2 = op(e4, e1)) | ~ sP71), inference(usedef, [], [e318])).
fof(e318, plain, (sP71 <=> ((~ (e2 = unit) & (e1 = op(e4, e2))) | ~ (e2 = op(e4, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP71])])).
fof(f319, plain, ((~ (e2 = unit) & (e2 = op(e4, e2))) | ~ (e2 = op(e4, e2)) | ~ sP72), inference(usedef, [], [e319])).
fof(e319, plain, (sP72 <=> ((~ (e2 = unit) & (e2 = op(e4, e2))) | ~ (e2 = op(e4, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP72])])).
fof(f320, plain, ((~ (e2 = unit) & (e3 = op(e4, e2))) | ~ (e2 = op(e4, e3)) | ~ sP73), inference(usedef, [], [e320])).
fof(e320, plain, (sP73 <=> ((~ (e2 = unit) & (e3 = op(e4, e2))) | ~ (e2 = op(e4, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP73])])).
fof(f321, plain, ((~ (e2 = unit) & (e4 = op(e4, e2))) | ~ (e2 = op(e4, e4)) | ~ sP74), inference(usedef, [], [e321])).
fof(e321, plain, (sP74 <=> ((~ (e2 = unit) & (e4 = op(e4, e2))) | ~ (e2 = op(e4, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP74])])).
fof(f322, plain, ((~ (e3 = unit) & (e0 = op(e0, e3))) | ~ (op(e0, e0) = e3) | ~ sP75), inference(usedef, [], [e322])).
fof(e322, plain, (sP75 <=> ((~ (e3 = unit) & (e0 = op(e0, e3))) | ~ (op(e0, e0) = e3))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP75])])).
fof(f323, plain, ((~ (e3 = unit) & (e1 = op(e0, e3))) | ~ (e3 = op(e0, e1)) | ~ sP76), inference(usedef, [], [e323])).
fof(e323, plain, (sP76 <=> ((~ (e3 = unit) & (e1 = op(e0, e3))) | ~ (e3 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP76])])).
fof(f324, plain, ((~ (e3 = unit) & (e2 = op(e0, e3))) | ~ (e3 = op(e0, e2)) | ~ sP77), inference(usedef, [], [e324])).
fof(e324, plain, (sP77 <=> ((~ (e3 = unit) & (e2 = op(e0, e3))) | ~ (e3 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP77])])).
fof(f325, plain, ((~ (e3 = unit) & (e3 = op(e0, e3))) | ~ (e3 = op(e0, e3)) | ~ sP78), inference(usedef, [], [e325])).
fof(e325, plain, (sP78 <=> ((~ (e3 = unit) & (e3 = op(e0, e3))) | ~ (e3 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP78])])).
fof(f326, plain, ((~ (e3 = unit) & (e4 = op(e0, e3))) | ~ (e3 = op(e0, e4)) | ~ sP79), inference(usedef, [], [e326])).
fof(e326, plain, (sP79 <=> ((~ (e3 = unit) & (e4 = op(e0, e3))) | ~ (e3 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP79])])).
fof(f327, plain, ((~ (e3 = unit) & (e0 = op(e1, e3))) | ~ (e3 = op(e1, e0)) | ~ sP80), inference(usedef, [], [e327])).
fof(e327, plain, (sP80 <=> ((~ (e3 = unit) & (e0 = op(e1, e3))) | ~ (e3 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP80])])).
fof(f328, plain, ((~ (e3 = unit) & (e1 = op(e1, e3))) | ~ (e3 = op(e1, e1)) | ~ sP81), inference(usedef, [], [e328])).
fof(e328, plain, (sP81 <=> ((~ (e3 = unit) & (e1 = op(e1, e3))) | ~ (e3 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP81])])).
fof(f329, plain, ((~ (e3 = unit) & (e2 = op(e1, e3))) | ~ (e3 = op(e1, e2)) | ~ sP82), inference(usedef, [], [e329])).
fof(e329, plain, (sP82 <=> ((~ (e3 = unit) & (e2 = op(e1, e3))) | ~ (e3 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP82])])).
fof(f330, plain, ((~ (e3 = unit) & (e3 = op(e1, e3))) | ~ (e3 = op(e1, e3)) | ~ sP83), inference(usedef, [], [e330])).
fof(e330, plain, (sP83 <=> ((~ (e3 = unit) & (e3 = op(e1, e3))) | ~ (e3 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP83])])).
fof(f331, plain, ((~ (e3 = unit) & (e4 = op(e1, e3))) | ~ (e3 = op(e1, e4)) | ~ sP84), inference(usedef, [], [e331])).
fof(e331, plain, (sP84 <=> ((~ (e3 = unit) & (e4 = op(e1, e3))) | ~ (e3 = op(e1, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP84])])).
fof(f332, plain, ((~ (e3 = unit) & (e0 = op(e2, e3))) | ~ (e3 = op(e2, e0)) | ~ sP85), inference(usedef, [], [e332])).
fof(e332, plain, (sP85 <=> ((~ (e3 = unit) & (e0 = op(e2, e3))) | ~ (e3 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP85])])).
fof(f333, plain, ((~ (e3 = unit) & (e1 = op(e2, e3))) | ~ (e3 = op(e2, e1)) | ~ sP86), inference(usedef, [], [e333])).
fof(e333, plain, (sP86 <=> ((~ (e3 = unit) & (e1 = op(e2, e3))) | ~ (e3 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP86])])).
fof(f334, plain, ((~ (e3 = unit) & (e2 = op(e2, e3))) | ~ (e3 = op(e2, e2)) | ~ sP87), inference(usedef, [], [e334])).
fof(e334, plain, (sP87 <=> ((~ (e3 = unit) & (e2 = op(e2, e3))) | ~ (e3 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP87])])).
fof(f335, plain, ((~ (e3 = unit) & (e3 = op(e2, e3))) | ~ (e3 = op(e2, e3)) | ~ sP88), inference(usedef, [], [e335])).
fof(e335, plain, (sP88 <=> ((~ (e3 = unit) & (e3 = op(e2, e3))) | ~ (e3 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP88])])).
fof(f336, plain, ((~ (e3 = unit) & (e4 = op(e2, e3))) | ~ (e3 = op(e2, e4)) | ~ sP89), inference(usedef, [], [e336])).
fof(e336, plain, (sP89 <=> ((~ (e3 = unit) & (e4 = op(e2, e3))) | ~ (e3 = op(e2, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP89])])).
fof(f337, plain, ((~ (e3 = unit) & (e0 = op(e3, e3))) | ~ (e3 = op(e3, e0)) | ~ sP90), inference(usedef, [], [e337])).
fof(e337, plain, (sP90 <=> ((~ (e3 = unit) & (e0 = op(e3, e3))) | ~ (e3 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP90])])).
fof(f338, plain, ((~ (e3 = unit) & (e1 = op(e3, e3))) | ~ (e3 = op(e3, e1)) | ~ sP91), inference(usedef, [], [e338])).
fof(e338, plain, (sP91 <=> ((~ (e3 = unit) & (e1 = op(e3, e3))) | ~ (e3 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP91])])).
fof(f339, plain, ((~ (e3 = unit) & (e2 = op(e3, e3))) | ~ (e3 = op(e3, e2)) | ~ sP92), inference(usedef, [], [e339])).
fof(e339, plain, (sP92 <=> ((~ (e3 = unit) & (e2 = op(e3, e3))) | ~ (e3 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP92])])).
fof(f340, plain, ((~ (e3 = unit) & (e3 = op(e3, e3))) | ~ (e3 = op(e3, e3)) | ~ sP93), inference(usedef, [], [e340])).
fof(e340, plain, (sP93 <=> ((~ (e3 = unit) & (e3 = op(e3, e3))) | ~ (e3 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP93])])).
fof(f341, plain, ((~ (e3 = unit) & (e4 = op(e3, e3))) | ~ (e3 = op(e3, e4)) | ~ sP94), inference(usedef, [], [e341])).
fof(e341, plain, (sP94 <=> ((~ (e3 = unit) & (e4 = op(e3, e3))) | ~ (e3 = op(e3, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP94])])).
fof(f342, plain, ((~ (e3 = unit) & (e0 = op(e4, e3))) | ~ (e3 = op(e4, e0)) | ~ sP95), inference(usedef, [], [e342])).
fof(e342, plain, (sP95 <=> ((~ (e3 = unit) & (e0 = op(e4, e3))) | ~ (e3 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP95])])).
fof(f343, plain, ((~ (e3 = unit) & (e1 = op(e4, e3))) | ~ (e3 = op(e4, e1)) | ~ sP96), inference(usedef, [], [e343])).
fof(e343, plain, (sP96 <=> ((~ (e3 = unit) & (e1 = op(e4, e3))) | ~ (e3 = op(e4, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP96])])).
fof(f344, plain, ((~ (e3 = unit) & (e2 = op(e4, e3))) | ~ (e3 = op(e4, e2)) | ~ sP97), inference(usedef, [], [e344])).
fof(e344, plain, (sP97 <=> ((~ (e3 = unit) & (e2 = op(e4, e3))) | ~ (e3 = op(e4, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP97])])).
fof(f345, plain, ((~ (e3 = unit) & (e3 = op(e4, e3))) | ~ (e3 = op(e4, e3)) | ~ sP98), inference(usedef, [], [e345])).
fof(e345, plain, (sP98 <=> ((~ (e3 = unit) & (e3 = op(e4, e3))) | ~ (e3 = op(e4, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP98])])).
fof(f346, plain, ((~ (e3 = unit) & (e4 = op(e4, e3))) | ~ (e3 = op(e4, e4)) | ~ sP99), inference(usedef, [], [e346])).
fof(e346, plain, (sP99 <=> ((~ (e3 = unit) & (e4 = op(e4, e3))) | ~ (e3 = op(e4, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP99])])).
fof(f347, plain, ((~ (e4 = unit) & (e0 = op(e0, e4))) | ~ (op(e0, e0) = e4) | ~ sP100), inference(usedef, [], [e347])).
fof(e347, plain, (sP100 <=> ((~ (e4 = unit) & (e0 = op(e0, e4))) | ~ (op(e0, e0) = e4))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP100])])).
fof(f348, plain, ((~ (e4 = unit) & (e1 = op(e0, e4))) | ~ (e4 = op(e0, e1)) | ~ sP101), inference(usedef, [], [e348])).
fof(e348, plain, (sP101 <=> ((~ (e4 = unit) & (e1 = op(e0, e4))) | ~ (e4 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP101])])).
fof(f349, plain, ((~ (e4 = unit) & (e2 = op(e0, e4))) | ~ (e4 = op(e0, e2)) | ~ sP102), inference(usedef, [], [e349])).
fof(e349, plain, (sP102 <=> ((~ (e4 = unit) & (e2 = op(e0, e4))) | ~ (e4 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP102])])).
fof(f350, plain, ((~ (e4 = unit) & (e3 = op(e0, e4))) | ~ (e4 = op(e0, e3)) | ~ sP103), inference(usedef, [], [e350])).
fof(e350, plain, (sP103 <=> ((~ (e4 = unit) & (e3 = op(e0, e4))) | ~ (e4 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP103])])).
fof(f351, plain, ((~ (e4 = unit) & (e4 = op(e0, e4))) | ~ (e4 = op(e0, e4)) | ~ sP104), inference(usedef, [], [e351])).
fof(e351, plain, (sP104 <=> ((~ (e4 = unit) & (e4 = op(e0, e4))) | ~ (e4 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP104])])).
fof(f352, plain, ((~ (e4 = unit) & (e0 = op(e1, e4))) | ~ (e4 = op(e1, e0)) | ~ sP105), inference(usedef, [], [e352])).
fof(e352, plain, (sP105 <=> ((~ (e4 = unit) & (e0 = op(e1, e4))) | ~ (e4 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP105])])).
fof(f353, plain, ((~ (e4 = unit) & (e1 = op(e1, e4))) | ~ (e4 = op(e1, e1)) | ~ sP106), inference(usedef, [], [e353])).
fof(e353, plain, (sP106 <=> ((~ (e4 = unit) & (e1 = op(e1, e4))) | ~ (e4 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP106])])).
fof(f354, plain, ((~ (e4 = unit) & (e2 = op(e1, e4))) | ~ (e4 = op(e1, e2)) | ~ sP107), inference(usedef, [], [e354])).
fof(e354, plain, (sP107 <=> ((~ (e4 = unit) & (e2 = op(e1, e4))) | ~ (e4 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP107])])).
fof(f355, plain, ((~ (e4 = unit) & (e3 = op(e1, e4))) | ~ (e4 = op(e1, e3)) | ~ sP108), inference(usedef, [], [e355])).
fof(e355, plain, (sP108 <=> ((~ (e4 = unit) & (e3 = op(e1, e4))) | ~ (e4 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP108])])).
fof(f356, plain, ((~ (e4 = unit) & (e4 = op(e1, e4))) | ~ (e4 = op(e1, e4)) | ~ sP109), inference(usedef, [], [e356])).
fof(e356, plain, (sP109 <=> ((~ (e4 = unit) & (e4 = op(e1, e4))) | ~ (e4 = op(e1, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP109])])).
fof(f357, plain, ((~ (e4 = unit) & (e0 = op(e2, e4))) | ~ (e4 = op(e2, e0)) | ~ sP110), inference(usedef, [], [e357])).
fof(e357, plain, (sP110 <=> ((~ (e4 = unit) & (e0 = op(e2, e4))) | ~ (e4 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP110])])).
fof(f358, plain, ((~ (e4 = unit) & (e1 = op(e2, e4))) | ~ (e4 = op(e2, e1)) | ~ sP111), inference(usedef, [], [e358])).
fof(e358, plain, (sP111 <=> ((~ (e4 = unit) & (e1 = op(e2, e4))) | ~ (e4 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP111])])).
fof(f359, plain, ((~ (e4 = unit) & (e2 = op(e2, e4))) | ~ (e4 = op(e2, e2)) | ~ sP112), inference(usedef, [], [e359])).
fof(e359, plain, (sP112 <=> ((~ (e4 = unit) & (e2 = op(e2, e4))) | ~ (e4 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP112])])).
fof(f360, plain, ((~ (e4 = unit) & (e3 = op(e2, e4))) | ~ (e4 = op(e2, e3)) | ~ sP113), inference(usedef, [], [e360])).
fof(e360, plain, (sP113 <=> ((~ (e4 = unit) & (e3 = op(e2, e4))) | ~ (e4 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP113])])).
fof(f361, plain, ((~ (e4 = unit) & (e4 = op(e2, e4))) | ~ (e4 = op(e2, e4)) | ~ sP114), inference(usedef, [], [e361])).
fof(e361, plain, (sP114 <=> ((~ (e4 = unit) & (e4 = op(e2, e4))) | ~ (e4 = op(e2, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP114])])).
fof(f362, plain, ((~ (e4 = unit) & (e0 = op(e3, e4))) | ~ (e4 = op(e3, e0)) | ~ sP115), inference(usedef, [], [e362])).
fof(e362, plain, (sP115 <=> ((~ (e4 = unit) & (e0 = op(e3, e4))) | ~ (e4 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP115])])).
fof(f363, plain, ((~ (e4 = unit) & (e1 = op(e3, e4))) | ~ (e4 = op(e3, e1)) | ~ sP116), inference(usedef, [], [e363])).
fof(e363, plain, (sP116 <=> ((~ (e4 = unit) & (e1 = op(e3, e4))) | ~ (e4 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP116])])).
fof(f364, plain, ((~ (e4 = unit) & (e2 = op(e3, e4))) | ~ (e4 = op(e3, e2)) | ~ sP117), inference(usedef, [], [e364])).
fof(e364, plain, (sP117 <=> ((~ (e4 = unit) & (e2 = op(e3, e4))) | ~ (e4 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP117])])).
fof(f365, plain, ((~ (e4 = unit) & (e3 = op(e3, e4))) | ~ (e4 = op(e3, e3)) | ~ sP118), inference(usedef, [], [e365])).
fof(e365, plain, (sP118 <=> ((~ (e4 = unit) & (e3 = op(e3, e4))) | ~ (e4 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP118])])).
fof(f366, plain, ((~ (e4 = unit) & (e4 = op(e3, e4))) | ~ (e4 = op(e3, e4)) | ~ sP119), inference(usedef, [], [e366])).
fof(e366, plain, (sP119 <=> ((~ (e4 = unit) & (e4 = op(e3, e4))) | ~ (e4 = op(e3, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP119])])).
fof(f367, plain, ((~ (e4 = unit) & (e0 = op(e4, e4))) | ~ (e4 = op(e4, e0)) | ~ sP120), inference(usedef, [], [e367])).
fof(e367, plain, (sP120 <=> ((~ (e4 = unit) & (e0 = op(e4, e4))) | ~ (e4 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP120])])).
fof(f368, plain, ((~ (e4 = unit) & (e1 = op(e4, e4))) | ~ (e4 = op(e4, e1)) | ~ sP121), inference(usedef, [], [e368])).
fof(e368, plain, (sP121 <=> ((~ (e4 = unit) & (e1 = op(e4, e4))) | ~ (e4 = op(e4, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP121])])).
fof(f369, plain, ((~ (e4 = unit) & (e2 = op(e4, e4))) | ~ (e4 = op(e4, e2)) | ~ sP122), inference(usedef, [], [e369])).
fof(e369, plain, (sP122 <=> ((~ (e4 = unit) & (e2 = op(e4, e4))) | ~ (e4 = op(e4, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP122])])).
fof(f370, plain, ((~ (e4 = unit) & (e3 = op(e4, e4))) | ~ (e4 = op(e4, e3)) | ~ sP123), inference(usedef, [], [e370])).
fof(e370, plain, (sP123 <=> ((~ (e4 = unit) & (e3 = op(e4, e4))) | ~ (e4 = op(e4, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP123])])).
fof(f371, plain, ((~ (e4 = unit) & (e4 = op(e4, e4))) | ~ (e4 = op(e4, e4)) | ~ sP124), inference(usedef, [], [e371])).
fof(e371, plain, (sP124 <=> ((~ (e4 = unit) & (e4 = op(e4, e4))) | ~ (e4 = op(e4, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP124])])).
fof(f372, plain, ((sP24 & sP23 & sP22 & sP21 & sP20 & sP19 & sP18 & sP17 & sP16 & sP15 & sP14 & sP13 & sP12 & sP11 & sP10 & sP9 & sP8 & sP7 & sP6 & sP5 & sP4 & sP3 & sP2 & sP1 & sP0) | ~ sP125), inference(usedef, [], [e372])).
fof(e372, plain, (sP125 <=> (sP24 & sP23 & sP22 & sP21 & sP20 & sP19 & sP18 & sP17 & sP16 & sP15 & sP14 & sP13 & sP12 & sP11 & sP10 & sP9 & sP8 & sP7 & sP6 & sP5 & sP4 & sP3 & sP2 & sP1 & sP0)), introduced(predicate_definition_introduction, [new_symbols(naming, [sP125])])).
fof(f373, plain, ((sP49 & sP48 & sP47 & sP46 & sP45 & sP44 & sP43 & sP42 & sP41 & sP40 & sP39 & sP38 & sP37 & sP36 & sP35 & sP34 & sP33 & sP32 & sP31 & sP30 & sP29 & sP28 & sP27 & sP26 & sP25) | ~ sP126), inference(usedef, [], [e373])).
fof(e373, plain, (sP126 <=> (sP49 & sP48 & sP47 & sP46 & sP45 & sP44 & sP43 & sP42 & sP41 & sP40 & sP39 & sP38 & sP37 & sP36 & sP35 & sP34 & sP33 & sP32 & sP31 & sP30 & sP29 & sP28 & sP27 & sP26 & sP25)), introduced(predicate_definition_introduction, [new_symbols(naming, [sP126])])).
fof(f374, plain, ((sP74 & sP73 & sP72 & sP71 & sP70 & sP69 & sP68 & sP67 & sP66 & sP65 & sP64 & sP63 & sP62 & sP61 & sP60 & sP59 & sP58 & sP57 & sP56 & sP55 & sP54 & sP53 & sP52 & sP51 & sP50) | ~ sP127), inference(usedef, [], [e374])).
fof(e374, plain, (sP127 <=> (sP74 & sP73 & sP72 & sP71 & sP70 & sP69 & sP68 & sP67 & sP66 & sP65 & sP64 & sP63 & sP62 & sP61 & sP60 & sP59 & sP58 & sP57 & sP56 & sP55 & sP54 & sP53 & sP52 & sP51 & sP50)), introduced(predicate_definition_introduction, [new_symbols(naming, [sP127])])).
fof(f375, plain, ((sP99 & sP98 & sP97 & sP96 & sP95 & sP94 & sP93 & sP92 & sP91 & sP90 & sP89 & sP88 & sP87 & sP86 & sP85 & sP84 & sP83 & sP82 & sP81 & sP80 & sP79 & sP78 & sP77 & sP76 & sP75) | ~ sP128), inference(usedef, [], [e375])).
fof(e375, plain, (sP128 <=> (sP99 & sP98 & sP97 & sP96 & sP95 & sP94 & sP93 & sP92 & sP91 & sP90 & sP89 & sP88 & sP87 & sP86 & sP85 & sP84 & sP83 & sP82 & sP81 & sP80 & sP79 & sP78 & sP77 & sP76 & sP75)), introduced(predicate_definition_introduction, [new_symbols(naming, [sP128])])).
fof(f376, plain, ((sP124 & sP123 & sP122 & sP121 & sP120 & sP119 & sP118 & sP117 & sP116 & sP115 & sP114 & sP113 & sP112 & sP111 & sP110 & sP109 & sP108 & sP107 & sP106 & sP105 & sP104 & sP103 & sP102 & sP101 & sP100) | ~ sP129), inference(usedef, [], [e376])).
fof(e376, plain, (sP129 <=> (sP124 & sP123 & sP122 & sP121 & sP120 & sP119 & sP118 & sP117 & sP116 & sP115 & sP114 & sP113 & sP112 & sP111 & sP110 & sP109 & sP108 & sP107 & sP106 & sP105 & sP104 & sP103 & sP102 & sP101 & sP100)), introduced(predicate_definition_introduction, [new_symbols(naming, [sP129])])).
fof(f6, plain, ((((~ (e4 = unit) & (e4 = op(e4, e4))) | ~ (e4 = op(e4, e4))) & ((~ (e4 = unit) & (e3 = op(e4, e4))) | ~ (e4 = op(e4, e3))) & ((~ (e4 = unit) & (e2 = op(e4, e4))) | ~ (e4 = op(e4, e2))) & ((~ (e4 = unit) & (e1 = op(e4, e4))) | ~ (e4 = op(e4, e1))) & ((~ (e4 = unit) & (e0 = op(e4, e4))) | ~ (e4 = op(e4, e0))) & ((~ (e4 = unit) & (e4 = op(e3, e4))) | ~ (e4 = op(e3, e4))) & ((~ (e4 = unit) & (e3 = op(e3, e4))) | ~ (e4 = op(e3, e3))) & ((~ (e4 = unit) & (e2 = op(e3, e4))) | ~ (e4 = op(e3, e2))) & ((~ (e4 = unit) & (e1 = op(e3, e4))) | ~ (e4 = op(e3, e1))) & ((~ (e4 = unit) & (e0 = op(e3, e4))) | ~ (e4 = op(e3, e0))) & ((~ (e4 = unit) & (e4 = op(e2, e4))) | ~ (e4 = op(e2, e4))) & ((~ (e4 = unit) & (e3 = op(e2, e4))) | ~ (e4 = op(e2, e3))) & ((~ (e4 = unit) & (e2 = op(e2, e4))) | ~ (e4 = op(e2, e2))) & ((~ (e4 = unit) & (e1 = op(e2, e4))) | ~ (e4 = op(e2, e1))) & ((~ (e4 = unit) & (e0 = op(e2, e4))) | ~ (e4 = op(e2, e0))) & ((~ (e4 = unit) & (e4 = op(e1, e4))) | ~ (e4 = op(e1, e4))) & ((~ (e4 = unit) & (e3 = op(e1, e4))) | ~ (e4 = op(e1, e3))) & ((~ (e4 = unit) & (e2 = op(e1, e4))) | ~ (e4 = op(e1, e2))) & ((~ (e4 = unit) & (e1 = op(e1, e4))) | ~ (e4 = op(e1, e1))) & ((~ (e4 = unit) & (e0 = op(e1, e4))) | ~ (e4 = op(e1, e0))) & ((~ (e4 = unit) & (e4 = op(e0, e4))) | ~ (e4 = op(e0, e4))) & ((~ (e4 = unit) & (e3 = op(e0, e4))) | ~ (e4 = op(e0, e3))) & ((~ (e4 = unit) & (e2 = op(e0, e4))) | ~ (e4 = op(e0, e2))) & ((~ (e4 = unit) & (e1 = op(e0, e4))) | ~ (e4 = op(e0, e1))) & ((~ (e4 = unit) & (e0 = op(e0, e4))) | ~ (op(e0, e0) = e4))) | (((~ (e3 = unit) & (e4 = op(e4, e3))) | ~ (e3 = op(e4, e4))) & ((~ (e3 = unit) & (e3 = op(e4, e3))) | ~ (e3 = op(e4, e3))) & ((~ (e3 = unit) & (e2 = op(e4, e3))) | ~ (e3 = op(e4, e2))) & ((~ (e3 = unit) & (e1 = op(e4, e3))) | ~ (e3 = op(e4, e1))) & ((~ (e3 = unit) & (e0 = op(e4, e3))) | ~ (e3 = op(e4, e0))) & ((~ (e3 = unit) & (e4 = op(e3, e3))) | ~ (e3 = op(e3, e4))) & ((~ (e3 = unit) & (e3 = op(e3, e3))) | ~ (e3 = op(e3, e3))) & ((~ (e3 = unit) & (e2 = op(e3, e3))) | ~ (e3 = op(e3, e2))) & ((~ (e3 = unit) & (e1 = op(e3, e3))) | ~ (e3 = op(e3, e1))) & ((~ (e3 = unit) & (e0 = op(e3, e3))) | ~ (e3 = op(e3, e0))) & ((~ (e3 = unit) & (e4 = op(e2, e3))) | ~ (e3 = op(e2, e4))) & ((~ (e3 = unit) & (e3 = op(e2, e3))) | ~ (e3 = op(e2, e3))) & ((~ (e3 = unit) & (e2 = op(e2, e3))) | ~ (e3 = op(e2, e2))) & ((~ (e3 = unit) & (e1 = op(e2, e3))) | ~ (e3 = op(e2, e1))) & ((~ (e3 = unit) & (e0 = op(e2, e3))) | ~ (e3 = op(e2, e0))) & ((~ (e3 = unit) & (e4 = op(e1, e3))) | ~ (e3 = op(e1, e4))) & ((~ (e3 = unit) & (e3 = op(e1, e3))) | ~ (e3 = op(e1, e3))) & ((~ (e3 = unit) & (e2 = op(e1, e3))) | ~ (e3 = op(e1, e2))) & ((~ (e3 = unit) & (e1 = op(e1, e3))) | ~ (e3 = op(e1, e1))) & ((~ (e3 = unit) & (e0 = op(e1, e3))) | ~ (e3 = op(e1, e0))) & ((~ (e3 = unit) & (e4 = op(e0, e3))) | ~ (e3 = op(e0, e4))) & ((~ (e3 = unit) & (e3 = op(e0, e3))) | ~ (e3 = op(e0, e3))) & ((~ (e3 = unit) & (e2 = op(e0, e3))) | ~ (e3 = op(e0, e2))) & ((~ (e3 = unit) & (e1 = op(e0, e3))) | ~ (e3 = op(e0, e1))) & ((~ (e3 = unit) & (e0 = op(e0, e3))) | ~ (op(e0, e0) = e3))) | (((~ (e2 = unit) & (e4 = op(e4, e2))) | ~ (e2 = op(e4, e4))) & ((~ (e2 = unit) & (e3 = op(e4, e2))) | ~ (e2 = op(e4, e3))) & ((~ (e2 = unit) & (e2 = op(e4, e2))) | ~ (e2 = op(e4, e2))) & ((~ (e2 = unit) & (e1 = op(e4, e2))) | ~ (e2 = op(e4, e1))) & ((~ (e2 = unit) & (e0 = op(e4, e2))) | ~ (e2 = op(e4, e0))) & ((~ (e2 = unit) & (e4 = op(e3, e2))) | ~ (e2 = op(e3, e4))) & ((~ (e2 = unit) & (e3 = op(e3, e2))) | ~ (e2 = op(e3, e3))) & ((~ (e2 = unit) & (e2 = op(e3, e2))) | ~ (e2 = op(e3, e2))) & ((~ (e2 = unit) & (e1 = op(e3, e2))) | ~ (e2 = op(e3, e1))) & ((~ (e2 = unit) & (e0 = op(e3, e2))) | ~ (e2 = op(e3, e0))) & ((~ (e2 = unit) & (e4 = op(e2, e2))) | ~ (e2 = op(e2, e4))) & ((~ (e2 = unit) & (e3 = op(e2, e2))) | ~ (e2 = op(e2, e3))) & ((~ (e2 = unit) & (e2 = op(e2, e2))) | ~ (e2 = op(e2, e2))) & ((~ (e2 = unit) & (e1 = op(e2, e2))) | ~ (e2 = op(e2, e1))) & ((~ (e2 = unit) & (e0 = op(e2, e2))) | ~ (e2 = op(e2, e0))) & ((~ (e2 = unit) & (e4 = op(e1, e2))) | ~ (e2 = op(e1, e4))) & ((~ (e2 = unit) & (e3 = op(e1, e2))) | ~ (e2 = op(e1, e3))) & ((~ (e2 = unit) & (e2 = op(e1, e2))) | ~ (e2 = op(e1, e2))) & ((~ (e2 = unit) & (e1 = op(e1, e2))) | ~ (e2 = op(e1, e1))) & ((~ (e2 = unit) & (e0 = op(e1, e2))) | ~ (e2 = op(e1, e0))) & ((~ (e2 = unit) & (e4 = op(e0, e2))) | ~ (e2 = op(e0, e4))) & ((~ (e2 = unit) & (e3 = op(e0, e2))) | ~ (e2 = op(e0, e3))) & ((~ (e2 = unit) & (e2 = op(e0, e2))) | ~ (e2 = op(e0, e2))) & ((~ (e2 = unit) & (e1 = op(e0, e2))) | ~ (e2 = op(e0, e1))) & ((~ (e2 = unit) & (e0 = op(e0, e2))) | ~ (op(e0, e0) = e2))) | (((~ (e1 = unit) & (e4 = op(e4, e1))) | ~ (e1 = op(e4, e4))) & ((~ (e1 = unit) & (e3 = op(e4, e1))) | ~ (e1 = op(e4, e3))) & ((~ (e1 = unit) & (e2 = op(e4, e1))) | ~ (e1 = op(e4, e2))) & ((~ (e1 = unit) & (e1 = op(e4, e1))) | ~ (e1 = op(e4, e1))) & ((~ (e1 = unit) & (e0 = op(e4, e1))) | ~ (e1 = op(e4, e0))) & ((~ (e1 = unit) & (e4 = op(e3, e1))) | ~ (e1 = op(e3, e4))) & ((~ (e1 = unit) & (e3 = op(e3, e1))) | ~ (e1 = op(e3, e3))) & ((~ (e1 = unit) & (e2 = op(e3, e1))) | ~ (e1 = op(e3, e2))) & ((~ (e1 = unit) & (e1 = op(e3, e1))) | ~ (e1 = op(e3, e1))) & ((~ (e1 = unit) & (e0 = op(e3, e1))) | ~ (e1 = op(e3, e0))) & ((~ (e1 = unit) & (e4 = op(e2, e1))) | ~ (e1 = op(e2, e4))) & ((~ (e1 = unit) & (e3 = op(e2, e1))) | ~ (e1 = op(e2, e3))) & ((~ (e1 = unit) & (e2 = op(e2, e1))) | ~ (e1 = op(e2, e2))) & ((~ (e1 = unit) & (e1 = op(e2, e1))) | ~ (e1 = op(e2, e1))) & ((~ (e1 = unit) & (e0 = op(e2, e1))) | ~ (e1 = op(e2, e0))) & ((~ (e1 = unit) & (e4 = op(e1, e1))) | ~ (e1 = op(e1, e4))) & ((~ (e1 = unit) & (e3 = op(e1, e1))) | ~ (e1 = op(e1, e3))) & ((~ (e1 = unit) & (e2 = op(e1, e1))) | ~ (e1 = op(e1, e2))) & ((~ (e1 = unit) & (e1 = op(e1, e1))) | ~ (e1 = op(e1, e1))) & ((~ (e1 = unit) & (e0 = op(e1, e1))) | ~ (e1 = op(e1, e0))) & ((~ (e1 = unit) & (e4 = op(e0, e1))) | ~ (e1 = op(e0, e4))) & ((~ (e1 = unit) & (e3 = op(e0, e1))) | ~ (e1 = op(e0, e3))) & ((~ (e1 = unit) & (e2 = op(e0, e1))) | ~ (e1 = op(e0, e2))) & ((~ (e1 = unit) & (e1 = op(e0, e1))) | ~ (e1 = op(e0, e1))) & ((~ (e1 = unit) & (e0 = op(e0, e1))) | ~ (op(e0, e0) = e1))) | (((~ (e0 = unit) & (e4 = op(e4, e0))) | ~ (e0 = op(e4, e4))) & ((~ (e0 = unit) & (e3 = op(e4, e0))) | ~ (e0 = op(e4, e3))) & ((~ (e0 = unit) & (e2 = op(e4, e0))) | ~ (e0 = op(e4, e2))) & ((~ (e0 = unit) & (e1 = op(e4, e0))) | ~ (e0 = op(e4, e1))) & ((~ (e0 = unit) & (e0 = op(e4, e0))) | ~ (e0 = op(e4, e0))) & ((~ (e0 = unit) & (e4 = op(e3, e0))) | ~ (e0 = op(e3, e4))) & ((~ (e0 = unit) & (e3 = op(e3, e0))) | ~ (e0 = op(e3, e3))) & ((~ (e0 = unit) & (e2 = op(e3, e0))) | ~ (e0 = op(e3, e2))) & ((~ (e0 = unit) & (e1 = op(e3, e0))) | ~ (e0 = op(e3, e1))) & ((~ (e0 = unit) & (e0 = op(e3, e0))) | ~ (e0 = op(e3, e0))) & ((~ (e0 = unit) & (e4 = op(e2, e0))) | ~ (e0 = op(e2, e4))) & ((~ (e0 = unit) & (e3 = op(e2, e0))) | ~ (e0 = op(e2, e3))) & ((~ (e0 = unit) & (e2 = op(e2, e0))) | ~ (e0 = op(e2, e2))) & ((~ (e0 = unit) & (e1 = op(e2, e0))) | ~ (e0 = op(e2, e1))) & ((~ (e0 = unit) & (e0 = op(e2, e0))) | ~ (e0 = op(e2, e0))) & ((~ (e0 = unit) & (e4 = op(e1, e0))) | ~ (e0 = op(e1, e4))) & ((~ (e0 = unit) & (e3 = op(e1, e0))) | ~ (e0 = op(e1, e3))) & ((~ (e0 = unit) & (e2 = op(e1, e0))) | ~ (e0 = op(e1, e2))) & ((~ (e0 = unit) & (e1 = op(e1, e0))) | ~ (e0 = op(e1, e1))) & ((~ (e0 = unit) & (e0 = op(e1, e0))) | ~ (e0 = op(e1, e0))) & ((~ (e0 = unit) & (op(e0, e0) = e4)) | ~ (e0 = op(e0, e4))) & ((~ (e0 = unit) & (op(e0, e0) = e3)) | ~ (e0 = op(e0, e3))) & ((~ (e0 = unit) & (op(e0, e0) = e2)) | ~ (e0 = op(e0, e2))) & ((~ (e0 = unit) & (op(e0, e0) = e1)) | ~ (e0 = op(e0, e1))) & ((~ (e0 = unit) & (e0 = op(e0, e0))) | ~ (e0 = op(e0, e0))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG099+1.p', ax6)).
fof(f2665, plain, (~ spl130_260 | ~ spl130_121 | ~ spl130_126), inference(avatar_split_clause, [], [f1078, f1726, f1705, f2437])).
fof(f2437, plain, (spl130_260 <=> sP0), introduced(avatar_definition, [new_symbols(naming, [spl130_260])])).
fof(f1078, plain, (~ (e0 = unit) | ~ (e0 = op(e0, e0)) | ~ sP0), inference(cnf_transformation, [], [f507])).
fof(f507, plain, ((~ (e0 = unit) & (e0 = op(e0, e0))) | ~ (e0 = op(e0, e0)) | ~ sP0), inference(nnf_transformation, [], [f247])).
fof(f2664, plain, (~ spl130_259 | ~ spl130_116 | spl130_122), inference(avatar_split_clause, [], [f1075, f1709, f1684, f2432])).
fof(f2432, plain, (spl130_259 <=> sP1), introduced(avatar_definition, [new_symbols(naming, [spl130_259])])).
fof(f1075, plain, ((op(e0, e0) = e1) | ~ (e0 = op(e0, e1)) | ~ sP1), inference(cnf_transformation, [], [f506])).
fof(f506, plain, ((~ (e0 = unit) & (op(e0, e0) = e1)) | ~ (e0 = op(e0, e1)) | ~ sP1), inference(nnf_transformation, [], [f248])).
fof(f2662, plain, (~ spl130_258 | ~ spl130_111 | spl130_123), inference(avatar_split_clause, [], [f1073, f1713, f1663, f2427])).
fof(f2427, plain, (spl130_258 <=> sP2), introduced(avatar_definition, [new_symbols(naming, [spl130_258])])).
fof(f1073, plain, ((op(e0, e0) = e2) | ~ (e0 = op(e0, e2)) | ~ sP2), inference(cnf_transformation, [], [f505])).
fof(f505, plain, ((~ (e0 = unit) & (op(e0, e0) = e2)) | ~ (e0 = op(e0, e2)) | ~ sP2), inference(nnf_transformation, [], [f249])).
fof(f2660, plain, (~ spl130_257 | ~ spl130_106 | spl130_124), inference(avatar_split_clause, [], [f1071, f1717, f1642, f2422])).
fof(f2422, plain, (spl130_257 <=> sP3), introduced(avatar_definition, [new_symbols(naming, [spl130_257])])).
fof(f1071, plain, ((op(e0, e0) = e3) | ~ (e0 = op(e0, e3)) | ~ sP3), inference(cnf_transformation, [], [f504])).
fof(f504, plain, ((~ (e0 = unit) & (op(e0, e0) = e3)) | ~ (e0 = op(e0, e3)) | ~ sP3), inference(nnf_transformation, [], [f250])).
fof(f2658, plain, (~ spl130_256 | ~ spl130_101 | spl130_125), inference(avatar_split_clause, [], [f1069, f1721, f1621, f2417])).
fof(f2417, plain, (spl130_256 <=> sP4), introduced(avatar_definition, [new_symbols(naming, [spl130_256])])).
fof(f1069, plain, ((op(e0, e0) = e4) | ~ (e0 = op(e0, e4)) | ~ sP4), inference(cnf_transformation, [], [f503])).
fof(f503, plain, ((~ (e0 = unit) & (op(e0, e0) = e4)) | ~ (e0 = op(e0, e4)) | ~ sP4), inference(nnf_transformation, [], [f251])).
fof(f2655, plain, (~ spl130_254 | ~ spl130_91 | spl130_97), inference(avatar_split_clause, [], [f1065, f1604, f1579, f2407])).
fof(f2407, plain, (spl130_254 <=> sP6), introduced(avatar_definition, [new_symbols(naming, [spl130_254])])).
fof(f1065, plain, ((e1 = op(e1, e0)) | ~ (e0 = op(e1, e1)) | ~ sP6), inference(cnf_transformation, [], [f501])).
fof(f501, plain, ((~ (e0 = unit) & (e1 = op(e1, e0))) | ~ (e0 = op(e1, e1)) | ~ sP6), inference(nnf_transformation, [], [f253])).
fof(f2653, plain, (~ spl130_253 | ~ spl130_86 | spl130_98), inference(avatar_split_clause, [], [f1063, f1608, f1558, f2402])).
fof(f2402, plain, (spl130_253 <=> sP7), introduced(avatar_definition, [new_symbols(naming, [spl130_253])])).
fof(f1063, plain, ((e2 = op(e1, e0)) | ~ (e0 = op(e1, e2)) | ~ sP7), inference(cnf_transformation, [], [f500])).
fof(f500, plain, ((~ (e0 = unit) & (e2 = op(e1, e0))) | ~ (e0 = op(e1, e2)) | ~ sP7), inference(nnf_transformation, [], [f254])).
fof(f2651, plain, (~ spl130_252 | ~ spl130_81 | spl130_99), inference(avatar_split_clause, [], [f1061, f1612, f1537, f2397])).
fof(f2397, plain, (spl130_252 <=> sP8), introduced(avatar_definition, [new_symbols(naming, [spl130_252])])).
fof(f1061, plain, ((e3 = op(e1, e0)) | ~ (e0 = op(e1, e3)) | ~ sP8), inference(cnf_transformation, [], [f499])).
fof(f499, plain, ((~ (e0 = unit) & (e3 = op(e1, e0))) | ~ (e0 = op(e1, e3)) | ~ sP8), inference(nnf_transformation, [], [f255])).
fof(f2649, plain, (~ spl130_251 | ~ spl130_76 | spl130_100), inference(avatar_split_clause, [], [f1059, f1616, f1516, f2392])).
fof(f2392, plain, (spl130_251 <=> sP9), introduced(avatar_definition, [new_symbols(naming, [spl130_251])])).
fof(f1059, plain, ((e4 = op(e1, e0)) | ~ (e0 = op(e1, e4)) | ~ sP9), inference(cnf_transformation, [], [f498])).
fof(f498, plain, ((~ (e0 = unit) & (e4 = op(e1, e0))) | ~ (e0 = op(e1, e4)) | ~ sP9), inference(nnf_transformation, [], [f256])).
fof(f2646, plain, (~ spl130_249 | ~ spl130_66 | spl130_72), inference(avatar_split_clause, [], [f1055, f1499, f1474, f2382])).
fof(f2382, plain, (spl130_249 <=> sP11), introduced(avatar_definition, [new_symbols(naming, [spl130_249])])).
fof(f1055, plain, ((e1 = op(e2, e0)) | ~ (e0 = op(e2, e1)) | ~ sP11), inference(cnf_transformation, [], [f496])).
fof(f496, plain, ((~ (e0 = unit) & (e1 = op(e2, e0))) | ~ (e0 = op(e2, e1)) | ~ sP11), inference(nnf_transformation, [], [f258])).
fof(f2644, plain, (~ spl130_248 | ~ spl130_61 | spl130_73), inference(avatar_split_clause, [], [f1053, f1503, f1453, f2377])).
fof(f2377, plain, (spl130_248 <=> sP12), introduced(avatar_definition, [new_symbols(naming, [spl130_248])])).
fof(f1053, plain, ((e2 = op(e2, e0)) | ~ (e0 = op(e2, e2)) | ~ sP12), inference(cnf_transformation, [], [f495])).
fof(f495, plain, ((~ (e0 = unit) & (e2 = op(e2, e0))) | ~ (e0 = op(e2, e2)) | ~ sP12), inference(nnf_transformation, [], [f259])).
fof(f2642, plain, (~ spl130_247 | ~ spl130_56 | spl130_74), inference(avatar_split_clause, [], [f1051, f1507, f1432, f2372])).
fof(f2372, plain, (spl130_247 <=> sP13), introduced(avatar_definition, [new_symbols(naming, [spl130_247])])).
fof(f1051, plain, ((e3 = op(e2, e0)) | ~ (e0 = op(e2, e3)) | ~ sP13), inference(cnf_transformation, [], [f494])).
fof(f494, plain, ((~ (e0 = unit) & (e3 = op(e2, e0))) | ~ (e0 = op(e2, e3)) | ~ sP13), inference(nnf_transformation, [], [f260])).
fof(f2640, plain, (~ spl130_246 | ~ spl130_51 | spl130_75), inference(avatar_split_clause, [], [f1049, f1511, f1411, f2367])).
fof(f2367, plain, (spl130_246 <=> sP14), introduced(avatar_definition, [new_symbols(naming, [spl130_246])])).
fof(f1049, plain, ((e4 = op(e2, e0)) | ~ (e0 = op(e2, e4)) | ~ sP14), inference(cnf_transformation, [], [f493])).
fof(f493, plain, ((~ (e0 = unit) & (e4 = op(e2, e0))) | ~ (e0 = op(e2, e4)) | ~ sP14), inference(nnf_transformation, [], [f261])).
fof(f2637, plain, (~ spl130_244 | ~ spl130_41 | spl130_47), inference(avatar_split_clause, [], [f1045, f1394, f1369, f2357])).
fof(f2357, plain, (spl130_244 <=> sP16), introduced(avatar_definition, [new_symbols(naming, [spl130_244])])).
fof(f1045, plain, ((e1 = op(e3, e0)) | ~ (e0 = op(e3, e1)) | ~ sP16), inference(cnf_transformation, [], [f491])).
fof(f491, plain, ((~ (e0 = unit) & (e1 = op(e3, e0))) | ~ (e0 = op(e3, e1)) | ~ sP16), inference(nnf_transformation, [], [f263])).
fof(f2635, plain, (~ spl130_243 | ~ spl130_36 | spl130_48), inference(avatar_split_clause, [], [f1043, f1398, f1348, f2352])).
fof(f2352, plain, (spl130_243 <=> sP17), introduced(avatar_definition, [new_symbols(naming, [spl130_243])])).
fof(f1043, plain, ((e2 = op(e3, e0)) | ~ (e0 = op(e3, e2)) | ~ sP17), inference(cnf_transformation, [], [f490])).
fof(f490, plain, ((~ (e0 = unit) & (e2 = op(e3, e0))) | ~ (e0 = op(e3, e2)) | ~ sP17), inference(nnf_transformation, [], [f264])).
fof(f2633, plain, (~ spl130_242 | ~ spl130_31 | spl130_49), inference(avatar_split_clause, [], [f1041, f1402, f1327, f2347])).
fof(f2347, plain, (spl130_242 <=> sP18), introduced(avatar_definition, [new_symbols(naming, [spl130_242])])).
fof(f1041, plain, ((e3 = op(e3, e0)) | ~ (e0 = op(e3, e3)) | ~ sP18), inference(cnf_transformation, [], [f489])).
fof(f489, plain, ((~ (e0 = unit) & (e3 = op(e3, e0))) | ~ (e0 = op(e3, e3)) | ~ sP18), inference(nnf_transformation, [], [f265])).
fof(f2631, plain, (~ spl130_241 | ~ spl130_26 | spl130_50), inference(avatar_split_clause, [], [f1039, f1406, f1306, f2342])).
fof(f2342, plain, (spl130_241 <=> sP19), introduced(avatar_definition, [new_symbols(naming, [spl130_241])])).
fof(f1039, plain, ((e4 = op(e3, e0)) | ~ (e0 = op(e3, e4)) | ~ sP19), inference(cnf_transformation, [], [f488])).
fof(f488, plain, ((~ (e0 = unit) & (e4 = op(e3, e0))) | ~ (e0 = op(e3, e4)) | ~ sP19), inference(nnf_transformation, [], [f266])).
fof(f2628, plain, (~ spl130_239 | ~ spl130_16 | spl130_22), inference(avatar_split_clause, [], [f1035, f1289, f1264, f2332])).
fof(f2332, plain, (spl130_239 <=> sP21), introduced(avatar_definition, [new_symbols(naming, [spl130_239])])).
fof(f1035, plain, ((e1 = op(e4, e0)) | ~ (e0 = op(e4, e1)) | ~ sP21), inference(cnf_transformation, [], [f486])).
fof(f486, plain, ((~ (e0 = unit) & (e1 = op(e4, e0))) | ~ (e0 = op(e4, e1)) | ~ sP21), inference(nnf_transformation, [], [f268])).
fof(f2626, plain, (~ spl130_238 | ~ spl130_11 | spl130_23), inference(avatar_split_clause, [], [f1033, f1293, f1243, f2327])).
fof(f2327, plain, (spl130_238 <=> sP22), introduced(avatar_definition, [new_symbols(naming, [spl130_238])])).
fof(f1033, plain, ((e2 = op(e4, e0)) | ~ (e0 = op(e4, e2)) | ~ sP22), inference(cnf_transformation, [], [f485])).
fof(f485, plain, ((~ (e0 = unit) & (e2 = op(e4, e0))) | ~ (e0 = op(e4, e2)) | ~ sP22), inference(nnf_transformation, [], [f269])).
fof(f2624, plain, (~ spl130_237 | ~ spl130_6 | spl130_24), inference(avatar_split_clause, [], [f1031, f1297, f1222, f2322])).
fof(f2322, plain, (spl130_237 <=> sP23), introduced(avatar_definition, [new_symbols(naming, [spl130_237])])).
fof(f1031, plain, ((e3 = op(e4, e0)) | ~ (e0 = op(e4, e3)) | ~ sP23), inference(cnf_transformation, [], [f484])).
fof(f484, plain, ((~ (e0 = unit) & (e3 = op(e4, e0))) | ~ (e0 = op(e4, e3)) | ~ sP23), inference(nnf_transformation, [], [f270])).
fof(f2622, plain, (~ spl130_236 | ~ spl130_1 | spl130_25), inference(avatar_split_clause, [], [f1029, f1301, f1201, f2317])).
fof(f2317, plain, (spl130_236 <=> sP24), introduced(avatar_definition, [new_symbols(naming, [spl130_236])])).
fof(f1029, plain, ((e4 = op(e4, e0)) | ~ (e0 = op(e4, e4)) | ~ sP24), inference(cnf_transformation, [], [f483])).
fof(f483, plain, ((~ (e0 = unit) & (e4 = op(e4, e0))) | ~ (e0 = op(e4, e4)) | ~ sP24), inference(nnf_transformation, [], [f271])).
fof(f2620, plain, (~ spl130_234 | ~ spl130_122 | spl130_116), inference(avatar_split_clause, [], [f1027, f1684, f1709, f2308])).
fof(f2308, plain, (spl130_234 <=> sP25), introduced(avatar_definition, [new_symbols(naming, [spl130_234])])).
fof(f1027, plain, ((e0 = op(e0, e1)) | ~ (op(e0, e0) = e1) | ~ sP25), inference(cnf_transformation, [], [f482])).
fof(f482, plain, ((~ (e1 = unit) & (e0 = op(e0, e1))) | ~ (op(e0, e0) = e1) | ~ sP25), inference(nnf_transformation, [], [f272])).
fof(f2617, plain, (~ spl130_232 | ~ spl130_112 | spl130_118), inference(avatar_split_clause, [], [f1023, f1692, f1667, f2298])).
fof(f2298, plain, (spl130_232 <=> sP27), introduced(avatar_definition, [new_symbols(naming, [spl130_232])])).
fof(f1023, plain, ((e2 = op(e0, e1)) | ~ (e1 = op(e0, e2)) | ~ sP27), inference(cnf_transformation, [], [f480])).
fof(f480, plain, ((~ (e1 = unit) & (e2 = op(e0, e1))) | ~ (e1 = op(e0, e2)) | ~ sP27), inference(nnf_transformation, [], [f274])).
fof(f2615, plain, (~ spl130_231 | ~ spl130_107 | spl130_119), inference(avatar_split_clause, [], [f1021, f1696, f1646, f2293])).
fof(f2293, plain, (spl130_231 <=> sP28), introduced(avatar_definition, [new_symbols(naming, [spl130_231])])).
fof(f1021, plain, ((e3 = op(e0, e1)) | ~ (e1 = op(e0, e3)) | ~ sP28), inference(cnf_transformation, [], [f479])).
fof(f479, plain, ((~ (e1 = unit) & (e3 = op(e0, e1))) | ~ (e1 = op(e0, e3)) | ~ sP28), inference(nnf_transformation, [], [f275])).
fof(f2613, plain, (~ spl130_230 | ~ spl130_102 | spl130_120), inference(avatar_split_clause, [], [f1019, f1700, f1625, f2288])).
fof(f2288, plain, (spl130_230 <=> sP29), introduced(avatar_definition, [new_symbols(naming, [spl130_230])])).
fof(f1019, plain, ((e4 = op(e0, e1)) | ~ (e1 = op(e0, e4)) | ~ sP29), inference(cnf_transformation, [], [f478])).
fof(f478, plain, ((~ (e1 = unit) & (e4 = op(e0, e1))) | ~ (e1 = op(e0, e4)) | ~ sP29), inference(nnf_transformation, [], [f276])).
fof(f2611, plain, (~ spl130_229 | ~ spl130_97 | spl130_91), inference(avatar_split_clause, [], [f1017, f1579, f1604, f2283])).
fof(f2283, plain, (spl130_229 <=> sP30), introduced(avatar_definition, [new_symbols(naming, [spl130_229])])).
fof(f1017, plain, ((e0 = op(e1, e1)) | ~ (e1 = op(e1, e0)) | ~ sP30), inference(cnf_transformation, [], [f477])).
fof(f477, plain, ((~ (e1 = unit) & (e0 = op(e1, e1))) | ~ (e1 = op(e1, e0)) | ~ sP30), inference(nnf_transformation, [], [f277])).
fof(f2609, plain, (~ spl130_228 | ~ spl130_92 | ~ spl130_127), inference(avatar_split_clause, [], [f1016, f1730, f1583, f2278])).
fof(f2278, plain, (spl130_228 <=> sP31), introduced(avatar_definition, [new_symbols(naming, [spl130_228])])).
fof(f1016, plain, (~ (e1 = unit) | ~ (e1 = op(e1, e1)) | ~ sP31), inference(cnf_transformation, [], [f476])).
fof(f476, plain, ((~ (e1 = unit) & (e1 = op(e1, e1))) | ~ (e1 = op(e1, e1)) | ~ sP31), inference(nnf_transformation, [], [f278])).
fof(f2608, plain, (~ spl130_227 | ~ spl130_87 | spl130_93), inference(avatar_split_clause, [], [f1013, f1587, f1562, f2273])).
fof(f2273, plain, (spl130_227 <=> sP32), introduced(avatar_definition, [new_symbols(naming, [spl130_227])])).
fof(f1013, plain, ((e2 = op(e1, e1)) | ~ (e1 = op(e1, e2)) | ~ sP32), inference(cnf_transformation, [], [f475])).
fof(f475, plain, ((~ (e1 = unit) & (e2 = op(e1, e1))) | ~ (e1 = op(e1, e2)) | ~ sP32), inference(nnf_transformation, [], [f279])).
fof(f2606, plain, (~ spl130_226 | ~ spl130_82 | spl130_94), inference(avatar_split_clause, [], [f1011, f1591, f1541, f2268])).
fof(f2268, plain, (spl130_226 <=> sP33), introduced(avatar_definition, [new_symbols(naming, [spl130_226])])).
fof(f1011, plain, ((e3 = op(e1, e1)) | ~ (e1 = op(e1, e3)) | ~ sP33), inference(cnf_transformation, [], [f474])).
fof(f474, plain, ((~ (e1 = unit) & (e3 = op(e1, e1))) | ~ (e1 = op(e1, e3)) | ~ sP33), inference(nnf_transformation, [], [f280])).
fof(f2604, plain, (~ spl130_225 | ~ spl130_77 | spl130_95), inference(avatar_split_clause, [], [f1009, f1595, f1520, f2263])).
fof(f2263, plain, (spl130_225 <=> sP34), introduced(avatar_definition, [new_symbols(naming, [spl130_225])])).
fof(f1009, plain, ((e4 = op(e1, e1)) | ~ (e1 = op(e1, e4)) | ~ sP34), inference(cnf_transformation, [], [f473])).
fof(f473, plain, ((~ (e1 = unit) & (e4 = op(e1, e1))) | ~ (e1 = op(e1, e4)) | ~ sP34), inference(nnf_transformation, [], [f281])).
fof(f2602, plain, (~ spl130_224 | ~ spl130_72 | spl130_66), inference(avatar_split_clause, [], [f1007, f1474, f1499, f2258])).
fof(f2258, plain, (spl130_224 <=> sP35), introduced(avatar_definition, [new_symbols(naming, [spl130_224])])).
fof(f1007, plain, ((e0 = op(e2, e1)) | ~ (e1 = op(e2, e0)) | ~ sP35), inference(cnf_transformation, [], [f472])).
fof(f472, plain, ((~ (e1 = unit) & (e0 = op(e2, e1))) | ~ (e1 = op(e2, e0)) | ~ sP35), inference(nnf_transformation, [], [f282])).
fof(f2599, plain, (~ spl130_222 | ~ spl130_62 | spl130_68), inference(avatar_split_clause, [], [f1003, f1482, f1457, f2248])).
fof(f2248, plain, (spl130_222 <=> sP37), introduced(avatar_definition, [new_symbols(naming, [spl130_222])])).
fof(f1003, plain, ((e2 = op(e2, e1)) | ~ (e1 = op(e2, e2)) | ~ sP37), inference(cnf_transformation, [], [f470])).
fof(f470, plain, ((~ (e1 = unit) & (e2 = op(e2, e1))) | ~ (e1 = op(e2, e2)) | ~ sP37), inference(nnf_transformation, [], [f284])).
fof(f2597, plain, (~ spl130_221 | ~ spl130_57 | spl130_69), inference(avatar_split_clause, [], [f1001, f1486, f1436, f2243])).
fof(f2243, plain, (spl130_221 <=> sP38), introduced(avatar_definition, [new_symbols(naming, [spl130_221])])).
fof(f1001, plain, ((e3 = op(e2, e1)) | ~ (e1 = op(e2, e3)) | ~ sP38), inference(cnf_transformation, [], [f469])).
fof(f469, plain, ((~ (e1 = unit) & (e3 = op(e2, e1))) | ~ (e1 = op(e2, e3)) | ~ sP38), inference(nnf_transformation, [], [f285])).
fof(f2595, plain, (~ spl130_220 | ~ spl130_52 | spl130_70), inference(avatar_split_clause, [], [f999, f1490, f1415, f2238])).
fof(f2238, plain, (spl130_220 <=> sP39), introduced(avatar_definition, [new_symbols(naming, [spl130_220])])).
fof(f999, plain, ((e4 = op(e2, e1)) | ~ (e1 = op(e2, e4)) | ~ sP39), inference(cnf_transformation, [], [f468])).
fof(f468, plain, ((~ (e1 = unit) & (e4 = op(e2, e1))) | ~ (e1 = op(e2, e4)) | ~ sP39), inference(nnf_transformation, [], [f286])).
fof(f2593, plain, (~ spl130_219 | ~ spl130_47 | spl130_41), inference(avatar_split_clause, [], [f997, f1369, f1394, f2233])).
fof(f2233, plain, (spl130_219 <=> sP40), introduced(avatar_definition, [new_symbols(naming, [spl130_219])])).
fof(f997, plain, ((e0 = op(e3, e1)) | ~ (e1 = op(e3, e0)) | ~ sP40), inference(cnf_transformation, [], [f467])).
fof(f467, plain, ((~ (e1 = unit) & (e0 = op(e3, e1))) | ~ (e1 = op(e3, e0)) | ~ sP40), inference(nnf_transformation, [], [f287])).
fof(f2590, plain, (~ spl130_217 | ~ spl130_37 | spl130_43), inference(avatar_split_clause, [], [f993, f1377, f1352, f2223])).
fof(f2223, plain, (spl130_217 <=> sP42), introduced(avatar_definition, [new_symbols(naming, [spl130_217])])).
fof(f993, plain, ((e2 = op(e3, e1)) | ~ (e1 = op(e3, e2)) | ~ sP42), inference(cnf_transformation, [], [f465])).
fof(f465, plain, ((~ (e1 = unit) & (e2 = op(e3, e1))) | ~ (e1 = op(e3, e2)) | ~ sP42), inference(nnf_transformation, [], [f289])).
fof(f2588, plain, (~ spl130_216 | ~ spl130_32 | spl130_44), inference(avatar_split_clause, [], [f991, f1381, f1331, f2218])).
fof(f2218, plain, (spl130_216 <=> sP43), introduced(avatar_definition, [new_symbols(naming, [spl130_216])])).
fof(f991, plain, ((e3 = op(e3, e1)) | ~ (e1 = op(e3, e3)) | ~ sP43), inference(cnf_transformation, [], [f464])).
fof(f464, plain, ((~ (e1 = unit) & (e3 = op(e3, e1))) | ~ (e1 = op(e3, e3)) | ~ sP43), inference(nnf_transformation, [], [f290])).
fof(f2586, plain, (~ spl130_215 | ~ spl130_27 | spl130_45), inference(avatar_split_clause, [], [f989, f1385, f1310, f2213])).
fof(f2213, plain, (spl130_215 <=> sP44), introduced(avatar_definition, [new_symbols(naming, [spl130_215])])).
fof(f989, plain, ((e4 = op(e3, e1)) | ~ (e1 = op(e3, e4)) | ~ sP44), inference(cnf_transformation, [], [f463])).
fof(f463, plain, ((~ (e1 = unit) & (e4 = op(e3, e1))) | ~ (e1 = op(e3, e4)) | ~ sP44), inference(nnf_transformation, [], [f291])).
fof(f2579, plain, (~ spl130_211 | ~ spl130_7 | spl130_19), inference(avatar_split_clause, [], [f981, f1276, f1226, f2193])).
fof(f2193, plain, (spl130_211 <=> sP48), introduced(avatar_definition, [new_symbols(naming, [spl130_211])])).
fof(f981, plain, ((e3 = op(e4, e1)) | ~ (e1 = op(e4, e3)) | ~ sP48), inference(cnf_transformation, [], [f459])).
fof(f459, plain, ((~ (e1 = unit) & (e3 = op(e4, e1))) | ~ (e1 = op(e4, e3)) | ~ sP48), inference(nnf_transformation, [], [f295])).
fof(f2577, plain, (~ spl130_210 | ~ spl130_2 | spl130_20), inference(avatar_split_clause, [], [f979, f1280, f1205, f2188])).
fof(f2188, plain, (spl130_210 <=> sP49), introduced(avatar_definition, [new_symbols(naming, [spl130_210])])).
fof(f979, plain, ((e4 = op(e4, e1)) | ~ (e1 = op(e4, e4)) | ~ sP49), inference(cnf_transformation, [], [f458])).
fof(f458, plain, ((~ (e1 = unit) & (e4 = op(e4, e1))) | ~ (e1 = op(e4, e4)) | ~ sP49), inference(nnf_transformation, [], [f296])).
fof(f2575, plain, (~ spl130_208 | ~ spl130_123 | spl130_111), inference(avatar_split_clause, [], [f977, f1663, f1713, f2179])).
fof(f2179, plain, (spl130_208 <=> sP50), introduced(avatar_definition, [new_symbols(naming, [spl130_208])])).
fof(f977, plain, ((e0 = op(e0, e2)) | ~ (op(e0, e0) = e2) | ~ sP50), inference(cnf_transformation, [], [f457])).
fof(f457, plain, ((~ (e2 = unit) & (e0 = op(e0, e2))) | ~ (op(e0, e0) = e2) | ~ sP50), inference(nnf_transformation, [], [f297])).
fof(f2574, plain, (~ spl130_208 | ~ spl130_123 | ~ spl130_128), inference(avatar_split_clause, [], [f978, f1734, f1713, f2179])).
fof(f978, plain, (~ (e2 = unit) | ~ (op(e0, e0) = e2) | ~ sP50), inference(cnf_transformation, [], [f457])).
fof(f2573, plain, (~ spl130_207 | ~ spl130_118 | spl130_112), inference(avatar_split_clause, [], [f975, f1667, f1692, f2174])).
fof(f2174, plain, (spl130_207 <=> sP51), introduced(avatar_definition, [new_symbols(naming, [spl130_207])])).
fof(f975, plain, ((e1 = op(e0, e2)) | ~ (e2 = op(e0, e1)) | ~ sP51), inference(cnf_transformation, [], [f456])).
fof(f456, plain, ((~ (e2 = unit) & (e1 = op(e0, e2))) | ~ (e2 = op(e0, e1)) | ~ sP51), inference(nnf_transformation, [], [f298])).
fof(f2570, plain, (~ spl130_205 | ~ spl130_108 | spl130_114), inference(avatar_split_clause, [], [f971, f1675, f1650, f2164])).
fof(f2164, plain, (spl130_205 <=> sP53), introduced(avatar_definition, [new_symbols(naming, [spl130_205])])).
fof(f971, plain, ((e3 = op(e0, e2)) | ~ (e2 = op(e0, e3)) | ~ sP53), inference(cnf_transformation, [], [f454])).
fof(f454, plain, ((~ (e2 = unit) & (e3 = op(e0, e2))) | ~ (e2 = op(e0, e3)) | ~ sP53), inference(nnf_transformation, [], [f300])).
fof(f2568, plain, (~ spl130_204 | ~ spl130_103 | spl130_115), inference(avatar_split_clause, [], [f969, f1679, f1629, f2159])).
fof(f2159, plain, (spl130_204 <=> sP54), introduced(avatar_definition, [new_symbols(naming, [spl130_204])])).
fof(f969, plain, ((e4 = op(e0, e2)) | ~ (e2 = op(e0, e4)) | ~ sP54), inference(cnf_transformation, [], [f453])).
fof(f453, plain, ((~ (e2 = unit) & (e4 = op(e0, e2))) | ~ (e2 = op(e0, e4)) | ~ sP54), inference(nnf_transformation, [], [f301])).
fof(f2566, plain, (~ spl130_203 | ~ spl130_98 | spl130_86), inference(avatar_split_clause, [], [f967, f1558, f1608, f2154])).
fof(f2154, plain, (spl130_203 <=> sP55), introduced(avatar_definition, [new_symbols(naming, [spl130_203])])).
fof(f967, plain, ((e0 = op(e1, e2)) | ~ (e2 = op(e1, e0)) | ~ sP55), inference(cnf_transformation, [], [f452])).
fof(f452, plain, ((~ (e2 = unit) & (e0 = op(e1, e2))) | ~ (e2 = op(e1, e0)) | ~ sP55), inference(nnf_transformation, [], [f302])).
fof(f2564, plain, (~ spl130_202 | ~ spl130_93 | spl130_87), inference(avatar_split_clause, [], [f965, f1562, f1587, f2149])).
fof(f2149, plain, (spl130_202 <=> sP56), introduced(avatar_definition, [new_symbols(naming, [spl130_202])])).
fof(f965, plain, ((e1 = op(e1, e2)) | ~ (e2 = op(e1, e1)) | ~ sP56), inference(cnf_transformation, [], [f451])).
fof(f451, plain, ((~ (e2 = unit) & (e1 = op(e1, e2))) | ~ (e2 = op(e1, e1)) | ~ sP56), inference(nnf_transformation, [], [f303])).
fof(f2561, plain, (~ spl130_200 | ~ spl130_83 | spl130_89), inference(avatar_split_clause, [], [f961, f1570, f1545, f2139])).
fof(f2139, plain, (spl130_200 <=> sP58), introduced(avatar_definition, [new_symbols(naming, [spl130_200])])).
fof(f961, plain, ((e3 = op(e1, e2)) | ~ (e2 = op(e1, e3)) | ~ sP58), inference(cnf_transformation, [], [f449])).
fof(f449, plain, ((~ (e2 = unit) & (e3 = op(e1, e2))) | ~ (e2 = op(e1, e3)) | ~ sP58), inference(nnf_transformation, [], [f305])).
fof(f2559, plain, (~ spl130_199 | ~ spl130_78 | spl130_90), inference(avatar_split_clause, [], [f959, f1574, f1524, f2134])).
fof(f2134, plain, (spl130_199 <=> sP59), introduced(avatar_definition, [new_symbols(naming, [spl130_199])])).
fof(f959, plain, ((e4 = op(e1, e2)) | ~ (e2 = op(e1, e4)) | ~ sP59), inference(cnf_transformation, [], [f448])).
fof(f448, plain, ((~ (e2 = unit) & (e4 = op(e1, e2))) | ~ (e2 = op(e1, e4)) | ~ sP59), inference(nnf_transformation, [], [f306])).
fof(f2557, plain, (~ spl130_198 | ~ spl130_73 | spl130_61), inference(avatar_split_clause, [], [f957, f1453, f1503, f2129])).
fof(f2129, plain, (spl130_198 <=> sP60), introduced(avatar_definition, [new_symbols(naming, [spl130_198])])).
fof(f957, plain, ((e0 = op(e2, e2)) | ~ (e2 = op(e2, e0)) | ~ sP60), inference(cnf_transformation, [], [f447])).
fof(f447, plain, ((~ (e2 = unit) & (e0 = op(e2, e2))) | ~ (e2 = op(e2, e0)) | ~ sP60), inference(nnf_transformation, [], [f307])).
fof(f2555, plain, (~ spl130_197 | ~ spl130_68 | spl130_62), inference(avatar_split_clause, [], [f955, f1457, f1482, f2124])).
fof(f2124, plain, (spl130_197 <=> sP61), introduced(avatar_definition, [new_symbols(naming, [spl130_197])])).
fof(f955, plain, ((e1 = op(e2, e2)) | ~ (e2 = op(e2, e1)) | ~ sP61), inference(cnf_transformation, [], [f446])).
fof(f446, plain, ((~ (e2 = unit) & (e1 = op(e2, e2))) | ~ (e2 = op(e2, e1)) | ~ sP61), inference(nnf_transformation, [], [f308])).
fof(f2553, plain, (~ spl130_196 | ~ spl130_63 | ~ spl130_128), inference(avatar_split_clause, [], [f954, f1734, f1461, f2119])).
fof(f2119, plain, (spl130_196 <=> sP62), introduced(avatar_definition, [new_symbols(naming, [spl130_196])])).
fof(f954, plain, (~ (e2 = unit) | ~ (e2 = op(e2, e2)) | ~ sP62), inference(cnf_transformation, [], [f445])).
fof(f445, plain, ((~ (e2 = unit) & (e2 = op(e2, e2))) | ~ (e2 = op(e2, e2)) | ~ sP62), inference(nnf_transformation, [], [f309])).
fof(f2552, plain, (~ spl130_195 | ~ spl130_58 | spl130_64), inference(avatar_split_clause, [], [f951, f1465, f1440, f2114])).
fof(f2114, plain, (spl130_195 <=> sP63), introduced(avatar_definition, [new_symbols(naming, [spl130_195])])).
fof(f951, plain, ((e3 = op(e2, e2)) | ~ (e2 = op(e2, e3)) | ~ sP63), inference(cnf_transformation, [], [f444])).
fof(f444, plain, ((~ (e2 = unit) & (e3 = op(e2, e2))) | ~ (e2 = op(e2, e3)) | ~ sP63), inference(nnf_transformation, [], [f310])).
fof(f2550, plain, (~ spl130_194 | ~ spl130_53 | spl130_65), inference(avatar_split_clause, [], [f949, f1469, f1419, f2109])).
fof(f2109, plain, (spl130_194 <=> sP64), introduced(avatar_definition, [new_symbols(naming, [spl130_194])])).
fof(f949, plain, ((e4 = op(e2, e2)) | ~ (e2 = op(e2, e4)) | ~ sP64), inference(cnf_transformation, [], [f443])).
fof(f443, plain, ((~ (e2 = unit) & (e4 = op(e2, e2))) | ~ (e2 = op(e2, e4)) | ~ sP64), inference(nnf_transformation, [], [f311])).
fof(f2548, plain, (~ spl130_193 | ~ spl130_48 | spl130_36), inference(avatar_split_clause, [], [f947, f1348, f1398, f2104])).
fof(f2104, plain, (spl130_193 <=> sP65), introduced(avatar_definition, [new_symbols(naming, [spl130_193])])).
fof(f947, plain, ((e0 = op(e3, e2)) | ~ (e2 = op(e3, e0)) | ~ sP65), inference(cnf_transformation, [], [f442])).
fof(f442, plain, ((~ (e2 = unit) & (e0 = op(e3, e2))) | ~ (e2 = op(e3, e0)) | ~ sP65), inference(nnf_transformation, [], [f312])).
fof(f2543, plain, (~ spl130_190 | ~ spl130_33 | spl130_39), inference(avatar_split_clause, [], [f941, f1360, f1335, f2089])).
fof(f2089, plain, (spl130_190 <=> sP68), introduced(avatar_definition, [new_symbols(naming, [spl130_190])])).
fof(f941, plain, ((e3 = op(e3, e2)) | ~ (e2 = op(e3, e3)) | ~ sP68), inference(cnf_transformation, [], [f439])).
fof(f439, plain, ((~ (e2 = unit) & (e3 = op(e3, e2))) | ~ (e2 = op(e3, e3)) | ~ sP68), inference(nnf_transformation, [], [f315])).
fof(f2541, plain, (~ spl130_189 | ~ spl130_28 | spl130_40), inference(avatar_split_clause, [], [f939, f1364, f1314, f2084])).
fof(f2084, plain, (spl130_189 <=> sP69), introduced(avatar_definition, [new_symbols(naming, [spl130_189])])).
fof(f939, plain, ((e4 = op(e3, e2)) | ~ (e2 = op(e3, e4)) | ~ sP69), inference(cnf_transformation, [], [f438])).
fof(f438, plain, ((~ (e2 = unit) & (e4 = op(e3, e2))) | ~ (e2 = op(e3, e4)) | ~ sP69), inference(nnf_transformation, [], [f316])).
fof(f2539, plain, (~ spl130_188 | ~ spl130_23 | spl130_11), inference(avatar_split_clause, [], [f937, f1243, f1293, f2079])).
fof(f2079, plain, (spl130_188 <=> sP70), introduced(avatar_definition, [new_symbols(naming, [spl130_188])])).
fof(f937, plain, ((e0 = op(e4, e2)) | ~ (e2 = op(e4, e0)) | ~ sP70), inference(cnf_transformation, [], [f437])).
fof(f437, plain, ((~ (e2 = unit) & (e0 = op(e4, e2))) | ~ (e2 = op(e4, e0)) | ~ sP70), inference(nnf_transformation, [], [f317])).
fof(f2534, plain, (~ spl130_185 | ~ spl130_8 | spl130_14), inference(avatar_split_clause, [], [f931, f1255, f1230, f2064])).
fof(f2064, plain, (spl130_185 <=> sP73), introduced(avatar_definition, [new_symbols(naming, [spl130_185])])).
fof(f931, plain, ((e3 = op(e4, e2)) | ~ (e2 = op(e4, e3)) | ~ sP73), inference(cnf_transformation, [], [f434])).
fof(f434, plain, ((~ (e2 = unit) & (e3 = op(e4, e2))) | ~ (e2 = op(e4, e3)) | ~ sP73), inference(nnf_transformation, [], [f320])).
fof(f2532, plain, (~ spl130_184 | ~ spl130_3 | spl130_15), inference(avatar_split_clause, [], [f929, f1259, f1209, f2059])).
fof(f2059, plain, (spl130_184 <=> sP74), introduced(avatar_definition, [new_symbols(naming, [spl130_184])])).
fof(f929, plain, ((e4 = op(e4, e2)) | ~ (e2 = op(e4, e4)) | ~ sP74), inference(cnf_transformation, [], [f433])).
fof(f433, plain, ((~ (e2 = unit) & (e4 = op(e4, e2))) | ~ (e2 = op(e4, e4)) | ~ sP74), inference(nnf_transformation, [], [f321])).
fof(f2530, plain, (~ spl130_182 | ~ spl130_124 | spl130_106), inference(avatar_split_clause, [], [f927, f1642, f1717, f2050])).
fof(f2050, plain, (spl130_182 <=> sP75), introduced(avatar_definition, [new_symbols(naming, [spl130_182])])).
fof(f927, plain, ((e0 = op(e0, e3)) | ~ (op(e0, e0) = e3) | ~ sP75), inference(cnf_transformation, [], [f432])).
fof(f432, plain, ((~ (e3 = unit) & (e0 = op(e0, e3))) | ~ (op(e0, e0) = e3) | ~ sP75), inference(nnf_transformation, [], [f322])).
fof(f2529, plain, (~ spl130_182 | ~ spl130_124 | ~ spl130_129), inference(avatar_split_clause, [], [f928, f1738, f1717, f2050])).
fof(f928, plain, (~ (e3 = unit) | ~ (op(e0, e0) = e3) | ~ sP75), inference(cnf_transformation, [], [f432])).
fof(f2528, plain, (~ spl130_181 | ~ spl130_119 | spl130_107), inference(avatar_split_clause, [], [f925, f1646, f1696, f2045])).
fof(f2045, plain, (spl130_181 <=> sP76), introduced(avatar_definition, [new_symbols(naming, [spl130_181])])).
fof(f925, plain, ((e1 = op(e0, e3)) | ~ (e3 = op(e0, e1)) | ~ sP76), inference(cnf_transformation, [], [f431])).
fof(f431, plain, ((~ (e3 = unit) & (e1 = op(e0, e3))) | ~ (e3 = op(e0, e1)) | ~ sP76), inference(nnf_transformation, [], [f323])).
fof(f2526, plain, (~ spl130_180 | ~ spl130_114 | spl130_108), inference(avatar_split_clause, [], [f923, f1650, f1675, f2040])).
fof(f2040, plain, (spl130_180 <=> sP77), introduced(avatar_definition, [new_symbols(naming, [spl130_180])])).
fof(f923, plain, ((e2 = op(e0, e3)) | ~ (e3 = op(e0, e2)) | ~ sP77), inference(cnf_transformation, [], [f430])).
fof(f430, plain, ((~ (e3 = unit) & (e2 = op(e0, e3))) | ~ (e3 = op(e0, e2)) | ~ sP77), inference(nnf_transformation, [], [f324])).
fof(f2523, plain, (~ spl130_178 | ~ spl130_104 | spl130_110), inference(avatar_split_clause, [], [f919, f1658, f1633, f2030])).
fof(f2030, plain, (spl130_178 <=> sP79), introduced(avatar_definition, [new_symbols(naming, [spl130_178])])).
fof(f919, plain, ((e4 = op(e0, e3)) | ~ (e3 = op(e0, e4)) | ~ sP79), inference(cnf_transformation, [], [f428])).
fof(f428, plain, ((~ (e3 = unit) & (e4 = op(e0, e3))) | ~ (e3 = op(e0, e4)) | ~ sP79), inference(nnf_transformation, [], [f326])).
fof(f2521, plain, (~ spl130_177 | ~ spl130_99 | spl130_81), inference(avatar_split_clause, [], [f917, f1537, f1612, f2025])).
fof(f2025, plain, (spl130_177 <=> sP80), introduced(avatar_definition, [new_symbols(naming, [spl130_177])])).
fof(f917, plain, ((e0 = op(e1, e3)) | ~ (e3 = op(e1, e0)) | ~ sP80), inference(cnf_transformation, [], [f427])).
fof(f427, plain, ((~ (e3 = unit) & (e0 = op(e1, e3))) | ~ (e3 = op(e1, e0)) | ~ sP80), inference(nnf_transformation, [], [f327])).
fof(f2520, plain, (~ spl130_177 | ~ spl130_99 | ~ spl130_129), inference(avatar_split_clause, [], [f918, f1738, f1612, f2025])).
fof(f918, plain, (~ (e3 = unit) | ~ (e3 = op(e1, e0)) | ~ sP80), inference(cnf_transformation, [], [f427])).
fof(f2519, plain, (~ spl130_176 | ~ spl130_94 | spl130_82), inference(avatar_split_clause, [], [f915, f1541, f1591, f2020])).
fof(f2020, plain, (spl130_176 <=> sP81), introduced(avatar_definition, [new_symbols(naming, [spl130_176])])).
fof(f915, plain, ((e1 = op(e1, e3)) | ~ (e3 = op(e1, e1)) | ~ sP81), inference(cnf_transformation, [], [f426])).
fof(f426, plain, ((~ (e3 = unit) & (e1 = op(e1, e3))) | ~ (e3 = op(e1, e1)) | ~ sP81), inference(nnf_transformation, [], [f328])).
fof(f2517, plain, (~ spl130_175 | ~ spl130_89 | spl130_83), inference(avatar_split_clause, [], [f913, f1545, f1570, f2015])).
fof(f2015, plain, (spl130_175 <=> sP82), introduced(avatar_definition, [new_symbols(naming, [spl130_175])])).
fof(f913, plain, ((e2 = op(e1, e3)) | ~ (e3 = op(e1, e2)) | ~ sP82), inference(cnf_transformation, [], [f425])).
fof(f425, plain, ((~ (e3 = unit) & (e2 = op(e1, e3))) | ~ (e3 = op(e1, e2)) | ~ sP82), inference(nnf_transformation, [], [f329])).
fof(f2514, plain, (~ spl130_173 | ~ spl130_79 | spl130_85), inference(avatar_split_clause, [], [f909, f1553, f1528, f2005])).
fof(f2005, plain, (spl130_173 <=> sP84), introduced(avatar_definition, [new_symbols(naming, [spl130_173])])).
fof(f909, plain, ((e4 = op(e1, e3)) | ~ (e3 = op(e1, e4)) | ~ sP84), inference(cnf_transformation, [], [f423])).
fof(f423, plain, ((~ (e3 = unit) & (e4 = op(e1, e3))) | ~ (e3 = op(e1, e4)) | ~ sP84), inference(nnf_transformation, [], [f331])).
fof(f2512, plain, (~ spl130_172 | ~ spl130_74 | spl130_56), inference(avatar_split_clause, [], [f907, f1432, f1507, f2000])).
fof(f2000, plain, (spl130_172 <=> sP85), introduced(avatar_definition, [new_symbols(naming, [spl130_172])])).
fof(f907, plain, ((e0 = op(e2, e3)) | ~ (e3 = op(e2, e0)) | ~ sP85), inference(cnf_transformation, [], [f422])).
fof(f422, plain, ((~ (e3 = unit) & (e0 = op(e2, e3))) | ~ (e3 = op(e2, e0)) | ~ sP85), inference(nnf_transformation, [], [f332])).
fof(f2510, plain, (~ spl130_171 | ~ spl130_69 | spl130_57), inference(avatar_split_clause, [], [f905, f1436, f1486, f1995])).
fof(f1995, plain, (spl130_171 <=> sP86), introduced(avatar_definition, [new_symbols(naming, [spl130_171])])).
fof(f905, plain, ((e1 = op(e2, e3)) | ~ (e3 = op(e2, e1)) | ~ sP86), inference(cnf_transformation, [], [f421])).
fof(f421, plain, ((~ (e3 = unit) & (e1 = op(e2, e3))) | ~ (e3 = op(e2, e1)) | ~ sP86), inference(nnf_transformation, [], [f333])).
fof(f2508, plain, (~ spl130_170 | ~ spl130_64 | spl130_58), inference(avatar_split_clause, [], [f903, f1440, f1465, f1990])).
fof(f1990, plain, (spl130_170 <=> sP87), introduced(avatar_definition, [new_symbols(naming, [spl130_170])])).
fof(f903, plain, ((e2 = op(e2, e3)) | ~ (e3 = op(e2, e2)) | ~ sP87), inference(cnf_transformation, [], [f420])).
fof(f420, plain, ((~ (e3 = unit) & (e2 = op(e2, e3))) | ~ (e3 = op(e2, e2)) | ~ sP87), inference(nnf_transformation, [], [f334])).
fof(f2505, plain, (~ spl130_168 | ~ spl130_54 | spl130_60), inference(avatar_split_clause, [], [f899, f1448, f1423, f1980])).
fof(f1980, plain, (spl130_168 <=> sP89), introduced(avatar_definition, [new_symbols(naming, [spl130_168])])).
fof(f899, plain, ((e4 = op(e2, e3)) | ~ (e3 = op(e2, e4)) | ~ sP89), inference(cnf_transformation, [], [f418])).
fof(f418, plain, ((~ (e3 = unit) & (e4 = op(e2, e3))) | ~ (e3 = op(e2, e4)) | ~ sP89), inference(nnf_transformation, [], [f336])).
fof(f2503, plain, (~ spl130_167 | ~ spl130_49 | spl130_31), inference(avatar_split_clause, [], [f897, f1327, f1402, f1975])).
fof(f1975, plain, (spl130_167 <=> sP90), introduced(avatar_definition, [new_symbols(naming, [spl130_167])])).
fof(f897, plain, ((e0 = op(e3, e3)) | ~ (e3 = op(e3, e0)) | ~ sP90), inference(cnf_transformation, [], [f417])).
fof(f417, plain, ((~ (e3 = unit) & (e0 = op(e3, e3))) | ~ (e3 = op(e3, e0)) | ~ sP90), inference(nnf_transformation, [], [f337])).
fof(f2501, plain, (~ spl130_166 | ~ spl130_44 | spl130_32), inference(avatar_split_clause, [], [f895, f1331, f1381, f1970])).
fof(f1970, plain, (spl130_166 <=> sP91), introduced(avatar_definition, [new_symbols(naming, [spl130_166])])).
fof(f895, plain, ((e1 = op(e3, e3)) | ~ (e3 = op(e3, e1)) | ~ sP91), inference(cnf_transformation, [], [f416])).
fof(f416, plain, ((~ (e3 = unit) & (e1 = op(e3, e3))) | ~ (e3 = op(e3, e1)) | ~ sP91), inference(nnf_transformation, [], [f338])).
fof(f2499, plain, (~ spl130_165 | ~ spl130_39 | spl130_33), inference(avatar_split_clause, [], [f893, f1335, f1360, f1965])).
fof(f1965, plain, (spl130_165 <=> sP92), introduced(avatar_definition, [new_symbols(naming, [spl130_165])])).
fof(f893, plain, ((e2 = op(e3, e3)) | ~ (e3 = op(e3, e2)) | ~ sP92), inference(cnf_transformation, [], [f415])).
fof(f415, plain, ((~ (e3 = unit) & (e2 = op(e3, e3))) | ~ (e3 = op(e3, e2)) | ~ sP92), inference(nnf_transformation, [], [f339])).
fof(f2497, plain, (~ spl130_164 | ~ spl130_34 | ~ spl130_129), inference(avatar_split_clause, [], [f892, f1738, f1339, f1960])).
fof(f1960, plain, (spl130_164 <=> sP93), introduced(avatar_definition, [new_symbols(naming, [spl130_164])])).
fof(f892, plain, (~ (e3 = unit) | ~ (e3 = op(e3, e3)) | ~ sP93), inference(cnf_transformation, [], [f414])).
fof(f414, plain, ((~ (e3 = unit) & (e3 = op(e3, e3))) | ~ (e3 = op(e3, e3)) | ~ sP93), inference(nnf_transformation, [], [f340])).
fof(f2496, plain, (~ spl130_163 | ~ spl130_29 | spl130_35), inference(avatar_split_clause, [], [f889, f1343, f1318, f1955])).
fof(f1955, plain, (spl130_163 <=> sP94), introduced(avatar_definition, [new_symbols(naming, [spl130_163])])).
fof(f889, plain, ((e4 = op(e3, e3)) | ~ (e3 = op(e3, e4)) | ~ sP94), inference(cnf_transformation, [], [f413])).
fof(f413, plain, ((~ (e3 = unit) & (e4 = op(e3, e3))) | ~ (e3 = op(e3, e4)) | ~ sP94), inference(nnf_transformation, [], [f341])).
fof(f2494, plain, (~ spl130_162 | ~ spl130_24 | spl130_6), inference(avatar_split_clause, [], [f887, f1222, f1297, f1950])).
fof(f1950, plain, (spl130_162 <=> sP95), introduced(avatar_definition, [new_symbols(naming, [spl130_162])])).
fof(f887, plain, ((e0 = op(e4, e3)) | ~ (e3 = op(e4, e0)) | ~ sP95), inference(cnf_transformation, [], [f412])).
fof(f412, plain, ((~ (e3 = unit) & (e0 = op(e4, e3))) | ~ (e3 = op(e4, e0)) | ~ sP95), inference(nnf_transformation, [], [f342])).
fof(f2490, plain, (~ spl130_160 | ~ spl130_14 | spl130_8), inference(avatar_split_clause, [], [f883, f1230, f1255, f1940])).
fof(f1940, plain, (spl130_160 <=> sP97), introduced(avatar_definition, [new_symbols(naming, [spl130_160])])).
fof(f883, plain, ((e2 = op(e4, e3)) | ~ (e3 = op(e4, e2)) | ~ sP97), inference(cnf_transformation, [], [f410])).
fof(f410, plain, ((~ (e3 = unit) & (e2 = op(e4, e3))) | ~ (e3 = op(e4, e2)) | ~ sP97), inference(nnf_transformation, [], [f344])).
fof(f2487, plain, (~ spl130_158 | ~ spl130_4 | spl130_10), inference(avatar_split_clause, [], [f879, f1238, f1213, f1930])).
fof(f1930, plain, (spl130_158 <=> sP99), introduced(avatar_definition, [new_symbols(naming, [spl130_158])])).
fof(f879, plain, ((e4 = op(e4, e3)) | ~ (e3 = op(e4, e4)) | ~ sP99), inference(cnf_transformation, [], [f408])).
fof(f408, plain, ((~ (e3 = unit) & (e4 = op(e4, e3))) | ~ (e3 = op(e4, e4)) | ~ sP99), inference(nnf_transformation, [], [f346])).
fof(f2485, plain, (~ spl130_156 | ~ spl130_125 | spl130_101), inference(avatar_split_clause, [], [f877, f1621, f1721, f1921])).
fof(f1921, plain, (spl130_156 <=> sP100), introduced(avatar_definition, [new_symbols(naming, [spl130_156])])).
fof(f877, plain, ((e0 = op(e0, e4)) | ~ (op(e0, e0) = e4) | ~ sP100), inference(cnf_transformation, [], [f407])).
fof(f407, plain, ((~ (e4 = unit) & (e0 = op(e0, e4))) | ~ (op(e0, e0) = e4) | ~ sP100), inference(nnf_transformation, [], [f347])).
fof(f2483, plain, (~ spl130_155 | ~ spl130_120 | spl130_102), inference(avatar_split_clause, [], [f875, f1625, f1700, f1916])).
fof(f1916, plain, (spl130_155 <=> sP101), introduced(avatar_definition, [new_symbols(naming, [spl130_155])])).
fof(f875, plain, ((e1 = op(e0, e4)) | ~ (e4 = op(e0, e1)) | ~ sP101), inference(cnf_transformation, [], [f406])).
fof(f406, plain, ((~ (e4 = unit) & (e1 = op(e0, e4))) | ~ (e4 = op(e0, e1)) | ~ sP101), inference(nnf_transformation, [], [f348])).
fof(f2481, plain, (~ spl130_154 | ~ spl130_115 | spl130_103), inference(avatar_split_clause, [], [f873, f1629, f1679, f1911])).
fof(f1911, plain, (spl130_154 <=> sP102), introduced(avatar_definition, [new_symbols(naming, [spl130_154])])).
fof(f873, plain, ((e2 = op(e0, e4)) | ~ (e4 = op(e0, e2)) | ~ sP102), inference(cnf_transformation, [], [f405])).
fof(f405, plain, ((~ (e4 = unit) & (e2 = op(e0, e4))) | ~ (e4 = op(e0, e2)) | ~ sP102), inference(nnf_transformation, [], [f349])).
fof(f2479, plain, (~ spl130_153 | ~ spl130_110 | spl130_104), inference(avatar_split_clause, [], [f871, f1633, f1658, f1906])).
fof(f1906, plain, (spl130_153 <=> sP103), introduced(avatar_definition, [new_symbols(naming, [spl130_153])])).
fof(f871, plain, ((e3 = op(e0, e4)) | ~ (e4 = op(e0, e3)) | ~ sP103), inference(cnf_transformation, [], [f404])).
fof(f404, plain, ((~ (e4 = unit) & (e3 = op(e0, e4))) | ~ (e4 = op(e0, e3)) | ~ sP103), inference(nnf_transformation, [], [f350])).
fof(f2476, plain, (~ spl130_151 | ~ spl130_100 | spl130_76), inference(avatar_split_clause, [], [f867, f1516, f1616, f1896])).
fof(f1896, plain, (spl130_151 <=> sP105), introduced(avatar_definition, [new_symbols(naming, [spl130_151])])).
fof(f867, plain, ((e0 = op(e1, e4)) | ~ (e4 = op(e1, e0)) | ~ sP105), inference(cnf_transformation, [], [f402])).
fof(f402, plain, ((~ (e4 = unit) & (e0 = op(e1, e4))) | ~ (e4 = op(e1, e0)) | ~ sP105), inference(nnf_transformation, [], [f352])).
fof(f2474, plain, (~ spl130_150 | ~ spl130_95 | spl130_77), inference(avatar_split_clause, [], [f865, f1520, f1595, f1891])).
fof(f1891, plain, (spl130_150 <=> sP106), introduced(avatar_definition, [new_symbols(naming, [spl130_150])])).
fof(f865, plain, ((e1 = op(e1, e4)) | ~ (e4 = op(e1, e1)) | ~ sP106), inference(cnf_transformation, [], [f401])).
fof(f401, plain, ((~ (e4 = unit) & (e1 = op(e1, e4))) | ~ (e4 = op(e1, e1)) | ~ sP106), inference(nnf_transformation, [], [f353])).
fof(f2472, plain, (~ spl130_149 | ~ spl130_90 | spl130_78), inference(avatar_split_clause, [], [f863, f1524, f1574, f1886])).
fof(f1886, plain, (spl130_149 <=> sP107), introduced(avatar_definition, [new_symbols(naming, [spl130_149])])).
fof(f863, plain, ((e2 = op(e1, e4)) | ~ (e4 = op(e1, e2)) | ~ sP107), inference(cnf_transformation, [], [f400])).
fof(f400, plain, ((~ (e4 = unit) & (e2 = op(e1, e4))) | ~ (e4 = op(e1, e2)) | ~ sP107), inference(nnf_transformation, [], [f354])).
fof(f2470, plain, (~ spl130_148 | ~ spl130_85 | spl130_79), inference(avatar_split_clause, [], [f861, f1528, f1553, f1881])).
fof(f1881, plain, (spl130_148 <=> sP108), introduced(avatar_definition, [new_symbols(naming, [spl130_148])])).
fof(f861, plain, ((e3 = op(e1, e4)) | ~ (e4 = op(e1, e3)) | ~ sP108), inference(cnf_transformation, [], [f399])).
fof(f399, plain, ((~ (e4 = unit) & (e3 = op(e1, e4))) | ~ (e4 = op(e1, e3)) | ~ sP108), inference(nnf_transformation, [], [f355])).
fof(f2467, plain, (~ spl130_146 | ~ spl130_75 | spl130_51), inference(avatar_split_clause, [], [f857, f1411, f1511, f1871])).
fof(f1871, plain, (spl130_146 <=> sP110), introduced(avatar_definition, [new_symbols(naming, [spl130_146])])).
fof(f857, plain, ((e0 = op(e2, e4)) | ~ (e4 = op(e2, e0)) | ~ sP110), inference(cnf_transformation, [], [f397])).
fof(f397, plain, ((~ (e4 = unit) & (e0 = op(e2, e4))) | ~ (e4 = op(e2, e0)) | ~ sP110), inference(nnf_transformation, [], [f357])).
fof(f2465, plain, (~ spl130_145 | ~ spl130_70 | spl130_52), inference(avatar_split_clause, [], [f855, f1415, f1490, f1866])).
fof(f1866, plain, (spl130_145 <=> sP111), introduced(avatar_definition, [new_symbols(naming, [spl130_145])])).
fof(f855, plain, ((e1 = op(e2, e4)) | ~ (e4 = op(e2, e1)) | ~ sP111), inference(cnf_transformation, [], [f396])).
fof(f396, plain, ((~ (e4 = unit) & (e1 = op(e2, e4))) | ~ (e4 = op(e2, e1)) | ~ sP111), inference(nnf_transformation, [], [f358])).
fof(f2463, plain, (~ spl130_144 | ~ spl130_65 | spl130_53), inference(avatar_split_clause, [], [f853, f1419, f1469, f1861])).
fof(f1861, plain, (spl130_144 <=> sP112), introduced(avatar_definition, [new_symbols(naming, [spl130_144])])).
fof(f853, plain, ((e2 = op(e2, e4)) | ~ (e4 = op(e2, e2)) | ~ sP112), inference(cnf_transformation, [], [f395])).
fof(f395, plain, ((~ (e4 = unit) & (e2 = op(e2, e4))) | ~ (e4 = op(e2, e2)) | ~ sP112), inference(nnf_transformation, [], [f359])).
fof(f2461, plain, (~ spl130_143 | ~ spl130_60 | spl130_54), inference(avatar_split_clause, [], [f851, f1423, f1448, f1856])).
fof(f1856, plain, (spl130_143 <=> sP113), introduced(avatar_definition, [new_symbols(naming, [spl130_143])])).
fof(f851, plain, ((e3 = op(e2, e4)) | ~ (e4 = op(e2, e3)) | ~ sP113), inference(cnf_transformation, [], [f394])).
fof(f394, plain, ((~ (e4 = unit) & (e3 = op(e2, e4))) | ~ (e4 = op(e2, e3)) | ~ sP113), inference(nnf_transformation, [], [f360])).
fof(f2456, plain, (~ spl130_140 | ~ spl130_45 | spl130_27), inference(avatar_split_clause, [], [f845, f1310, f1385, f1841])).
fof(f1841, plain, (spl130_140 <=> sP116), introduced(avatar_definition, [new_symbols(naming, [spl130_140])])).
fof(f845, plain, ((e1 = op(e3, e4)) | ~ (e4 = op(e3, e1)) | ~ sP116), inference(cnf_transformation, [], [f391])).
fof(f391, plain, ((~ (e4 = unit) & (e1 = op(e3, e4))) | ~ (e4 = op(e3, e1)) | ~ sP116), inference(nnf_transformation, [], [f363])).
fof(f2454, plain, (~ spl130_139 | ~ spl130_40 | spl130_28), inference(avatar_split_clause, [], [f843, f1314, f1364, f1836])).
fof(f1836, plain, (spl130_139 <=> sP117), introduced(avatar_definition, [new_symbols(naming, [spl130_139])])).
fof(f843, plain, ((e2 = op(e3, e4)) | ~ (e4 = op(e3, e2)) | ~ sP117), inference(cnf_transformation, [], [f390])).
fof(f390, plain, ((~ (e4 = unit) & (e2 = op(e3, e4))) | ~ (e4 = op(e3, e2)) | ~ sP117), inference(nnf_transformation, [], [f364])).
fof(f2452, plain, (~ spl130_138 | ~ spl130_35 | spl130_29), inference(avatar_split_clause, [], [f841, f1318, f1343, f1831])).
fof(f1831, plain, (spl130_138 <=> sP118), introduced(avatar_definition, [new_symbols(naming, [spl130_138])])).
fof(f841, plain, ((e3 = op(e3, e4)) | ~ (e4 = op(e3, e3)) | ~ sP118), inference(cnf_transformation, [], [f389])).
fof(f389, plain, ((~ (e4 = unit) & (e3 = op(e3, e4))) | ~ (e4 = op(e3, e3)) | ~ sP118), inference(nnf_transformation, [], [f365])).
fof(f2449, plain, (~ spl130_136 | ~ spl130_25 | spl130_1), inference(avatar_split_clause, [], [f837, f1201, f1301, f1821])).
fof(f1821, plain, (spl130_136 <=> sP120), introduced(avatar_definition, [new_symbols(naming, [spl130_136])])).
fof(f837, plain, ((e0 = op(e4, e4)) | ~ (e4 = op(e4, e0)) | ~ sP120), inference(cnf_transformation, [], [f387])).
fof(f387, plain, ((~ (e4 = unit) & (e0 = op(e4, e4))) | ~ (e4 = op(e4, e0)) | ~ sP120), inference(nnf_transformation, [], [f367])).
fof(f2447, plain, (~ spl130_135 | ~ spl130_20 | spl130_2), inference(avatar_split_clause, [], [f835, f1205, f1280, f1816])).
fof(f1816, plain, (spl130_135 <=> sP121), introduced(avatar_definition, [new_symbols(naming, [spl130_135])])).
fof(f835, plain, ((e1 = op(e4, e4)) | ~ (e4 = op(e4, e1)) | ~ sP121), inference(cnf_transformation, [], [f386])).
fof(f386, plain, ((~ (e4 = unit) & (e1 = op(e4, e4))) | ~ (e4 = op(e4, e1)) | ~ sP121), inference(nnf_transformation, [], [f368])).
fof(f2446, plain, (~ spl130_135 | ~ spl130_20 | ~ spl130_130), inference(avatar_split_clause, [], [f836, f1742, f1280, f1816])).
fof(f836, plain, (~ (e4 = unit) | ~ (e4 = op(e4, e1)) | ~ sP121), inference(cnf_transformation, [], [f386])).
fof(f2445, plain, (~ spl130_134 | ~ spl130_15 | spl130_3), inference(avatar_split_clause, [], [f833, f1209, f1259, f1811])).
fof(f1811, plain, (spl130_134 <=> sP122), introduced(avatar_definition, [new_symbols(naming, [spl130_134])])).
fof(f833, plain, ((e2 = op(e4, e4)) | ~ (e4 = op(e4, e2)) | ~ sP122), inference(cnf_transformation, [], [f385])).
fof(f385, plain, ((~ (e4 = unit) & (e2 = op(e4, e4))) | ~ (e4 = op(e4, e2)) | ~ sP122), inference(nnf_transformation, [], [f369])).
fof(f2443, plain, (~ spl130_133 | ~ spl130_10 | spl130_4), inference(avatar_split_clause, [], [f831, f1213, f1238, f1806])).
fof(f1806, plain, (spl130_133 <=> sP123), introduced(avatar_definition, [new_symbols(naming, [spl130_133])])).
fof(f831, plain, ((e3 = op(e4, e4)) | ~ (e4 = op(e4, e3)) | ~ sP123), inference(cnf_transformation, [], [f384])).
fof(f384, plain, ((~ (e4 = unit) & (e3 = op(e4, e4))) | ~ (e4 = op(e4, e3)) | ~ sP123), inference(nnf_transformation, [], [f370])).
fof(f2441, plain, (~ spl130_132 | ~ spl130_5 | ~ spl130_130), inference(avatar_split_clause, [], [f830, f1742, f1217, f1801])).
fof(f1801, plain, (spl130_132 <=> sP124), introduced(avatar_definition, [new_symbols(naming, [spl130_132])])).
fof(f830, plain, (~ (e4 = unit) | ~ (e4 = op(e4, e4)) | ~ sP124), inference(cnf_transformation, [], [f383])).
fof(f383, plain, ((~ (e4 = unit) & (e4 = op(e4, e4))) | ~ (e4 = op(e4, e4)) | ~ sP124), inference(nnf_transformation, [], [f371])).
fof(f2440, plain, (~ spl130_235 | spl130_260), inference(avatar_split_clause, [], [f804, f2437, f2313])).
fof(f804, plain, (sP0 | ~ sP125), inference(cnf_transformation, [], [f382])).
fof(f382, plain, ((sP24 & sP23 & sP22 & sP21 & sP20 & sP19 & sP18 & sP17 & sP16 & sP15 & sP14 & sP13 & sP12 & sP11 & sP10 & sP9 & sP8 & sP7 & sP6 & sP5 & sP4 & sP3 & sP2 & sP1 & sP0) | ~ sP125), inference(nnf_transformation, [], [f372])).
fof(f2435, plain, (~ spl130_235 | spl130_259), inference(avatar_split_clause, [], [f805, f2432, f2313])).
fof(f805, plain, (sP1 | ~ sP125), inference(cnf_transformation, [], [f382])).
fof(f2430, plain, (~ spl130_235 | spl130_258), inference(avatar_split_clause, [], [f806, f2427, f2313])).
fof(f806, plain, (sP2 | ~ sP125), inference(cnf_transformation, [], [f382])).
fof(f2425, plain, (~ spl130_235 | spl130_257), inference(avatar_split_clause, [], [f807, f2422, f2313])).
fof(f807, plain, (sP3 | ~ sP125), inference(cnf_transformation, [], [f382])).
fof(f2420, plain, (~ spl130_235 | spl130_256), inference(avatar_split_clause, [], [f808, f2417, f2313])).
fof(f808, plain, (sP4 | ~ sP125), inference(cnf_transformation, [], [f382])).
fof(f2410, plain, (~ spl130_235 | spl130_254), inference(avatar_split_clause, [], [f810, f2407, f2313])).
fof(f810, plain, (sP6 | ~ sP125), inference(cnf_transformation, [], [f382])).
fof(f2405, plain, (~ spl130_235 | spl130_253), inference(avatar_split_clause, [], [f811, f2402, f2313])).
fof(f811, plain, (sP7 | ~ sP125), inference(cnf_transformation, [], [f382])).
fof(f2400, plain, (~ spl130_235 | spl130_252), inference(avatar_split_clause, [], [f812, f2397, f2313])).
fof(f812, plain, (sP8 | ~ sP125), inference(cnf_transformation, [], [f382])).
fof(f2395, plain, (~ spl130_235 | spl130_251), inference(avatar_split_clause, [], [f813, f2392, f2313])).
fof(f813, plain, (sP9 | ~ sP125), inference(cnf_transformation, [], [f382])).
fof(f2385, plain, (~ spl130_235 | spl130_249), inference(avatar_split_clause, [], [f815, f2382, f2313])).
fof(f815, plain, (sP11 | ~ sP125), inference(cnf_transformation, [], [f382])).
fof(f2380, plain, (~ spl130_235 | spl130_248), inference(avatar_split_clause, [], [f816, f2377, f2313])).
fof(f816, plain, (sP12 | ~ sP125), inference(cnf_transformation, [], [f382])).
fof(f2375, plain, (~ spl130_235 | spl130_247), inference(avatar_split_clause, [], [f817, f2372, f2313])).
fof(f817, plain, (sP13 | ~ sP125), inference(cnf_transformation, [], [f382])).
fof(f2370, plain, (~ spl130_235 | spl130_246), inference(avatar_split_clause, [], [f818, f2367, f2313])).
fof(f818, plain, (sP14 | ~ sP125), inference(cnf_transformation, [], [f382])).
fof(f2360, plain, (~ spl130_235 | spl130_244), inference(avatar_split_clause, [], [f820, f2357, f2313])).
fof(f820, plain, (sP16 | ~ sP125), inference(cnf_transformation, [], [f382])).
fof(f2355, plain, (~ spl130_235 | spl130_243), inference(avatar_split_clause, [], [f821, f2352, f2313])).
fof(f821, plain, (sP17 | ~ sP125), inference(cnf_transformation, [], [f382])).
fof(f2350, plain, (~ spl130_235 | spl130_242), inference(avatar_split_clause, [], [f822, f2347, f2313])).
fof(f822, plain, (sP18 | ~ sP125), inference(cnf_transformation, [], [f382])).
fof(f2345, plain, (~ spl130_235 | spl130_241), inference(avatar_split_clause, [], [f823, f2342, f2313])).
fof(f823, plain, (sP19 | ~ sP125), inference(cnf_transformation, [], [f382])).
fof(f2335, plain, (~ spl130_235 | spl130_239), inference(avatar_split_clause, [], [f825, f2332, f2313])).
fof(f825, plain, (sP21 | ~ sP125), inference(cnf_transformation, [], [f382])).
fof(f2330, plain, (~ spl130_235 | spl130_238), inference(avatar_split_clause, [], [f826, f2327, f2313])).
fof(f826, plain, (sP22 | ~ sP125), inference(cnf_transformation, [], [f382])).
fof(f2325, plain, (~ spl130_235 | spl130_237), inference(avatar_split_clause, [], [f827, f2322, f2313])).
fof(f827, plain, (sP23 | ~ sP125), inference(cnf_transformation, [], [f382])).
fof(f2320, plain, (~ spl130_235 | spl130_236), inference(avatar_split_clause, [], [f828, f2317, f2313])).
fof(f828, plain, (sP24 | ~ sP125), inference(cnf_transformation, [], [f382])).
fof(f2311, plain, (~ spl130_209 | spl130_234), inference(avatar_split_clause, [], [f779, f2308, f2184])).
fof(f779, plain, (sP25 | ~ sP126), inference(cnf_transformation, [], [f381])).
fof(f381, plain, ((sP49 & sP48 & sP47 & sP46 & sP45 & sP44 & sP43 & sP42 & sP41 & sP40 & sP39 & sP38 & sP37 & sP36 & sP35 & sP34 & sP33 & sP32 & sP31 & sP30 & sP29 & sP28 & sP27 & sP26 & sP25) | ~ sP126), inference(nnf_transformation, [], [f373])).
fof(f2301, plain, (~ spl130_209 | spl130_232), inference(avatar_split_clause, [], [f781, f2298, f2184])).
fof(f781, plain, (sP27 | ~ sP126), inference(cnf_transformation, [], [f381])).
fof(f2296, plain, (~ spl130_209 | spl130_231), inference(avatar_split_clause, [], [f782, f2293, f2184])).
fof(f782, plain, (sP28 | ~ sP126), inference(cnf_transformation, [], [f381])).
fof(f2291, plain, (~ spl130_209 | spl130_230), inference(avatar_split_clause, [], [f783, f2288, f2184])).
fof(f783, plain, (sP29 | ~ sP126), inference(cnf_transformation, [], [f381])).
fof(f2286, plain, (~ spl130_209 | spl130_229), inference(avatar_split_clause, [], [f784, f2283, f2184])).
fof(f784, plain, (sP30 | ~ sP126), inference(cnf_transformation, [], [f381])).
fof(f2281, plain, (~ spl130_209 | spl130_228), inference(avatar_split_clause, [], [f785, f2278, f2184])).
fof(f785, plain, (sP31 | ~ sP126), inference(cnf_transformation, [], [f381])).
fof(f2276, plain, (~ spl130_209 | spl130_227), inference(avatar_split_clause, [], [f786, f2273, f2184])).
fof(f786, plain, (sP32 | ~ sP126), inference(cnf_transformation, [], [f381])).
fof(f2271, plain, (~ spl130_209 | spl130_226), inference(avatar_split_clause, [], [f787, f2268, f2184])).
fof(f787, plain, (sP33 | ~ sP126), inference(cnf_transformation, [], [f381])).
fof(f2266, plain, (~ spl130_209 | spl130_225), inference(avatar_split_clause, [], [f788, f2263, f2184])).
fof(f788, plain, (sP34 | ~ sP126), inference(cnf_transformation, [], [f381])).
fof(f2261, plain, (~ spl130_209 | spl130_224), inference(avatar_split_clause, [], [f789, f2258, f2184])).
fof(f789, plain, (sP35 | ~ sP126), inference(cnf_transformation, [], [f381])).
fof(f2251, plain, (~ spl130_209 | spl130_222), inference(avatar_split_clause, [], [f791, f2248, f2184])).
fof(f791, plain, (sP37 | ~ sP126), inference(cnf_transformation, [], [f381])).
fof(f2246, plain, (~ spl130_209 | spl130_221), inference(avatar_split_clause, [], [f792, f2243, f2184])).
fof(f792, plain, (sP38 | ~ sP126), inference(cnf_transformation, [], [f381])).
fof(f2241, plain, (~ spl130_209 | spl130_220), inference(avatar_split_clause, [], [f793, f2238, f2184])).
fof(f793, plain, (sP39 | ~ sP126), inference(cnf_transformation, [], [f381])).
fof(f2236, plain, (~ spl130_209 | spl130_219), inference(avatar_split_clause, [], [f794, f2233, f2184])).
fof(f794, plain, (sP40 | ~ sP126), inference(cnf_transformation, [], [f381])).
fof(f2226, plain, (~ spl130_209 | spl130_217), inference(avatar_split_clause, [], [f796, f2223, f2184])).
fof(f796, plain, (sP42 | ~ sP126), inference(cnf_transformation, [], [f381])).
fof(f2221, plain, (~ spl130_209 | spl130_216), inference(avatar_split_clause, [], [f797, f2218, f2184])).
fof(f797, plain, (sP43 | ~ sP126), inference(cnf_transformation, [], [f381])).
fof(f2216, plain, (~ spl130_209 | spl130_215), inference(avatar_split_clause, [], [f798, f2213, f2184])).
fof(f798, plain, (sP44 | ~ sP126), inference(cnf_transformation, [], [f381])).
fof(f2196, plain, (~ spl130_209 | spl130_211), inference(avatar_split_clause, [], [f802, f2193, f2184])).
fof(f802, plain, (sP48 | ~ sP126), inference(cnf_transformation, [], [f381])).
fof(f2191, plain, (~ spl130_209 | spl130_210), inference(avatar_split_clause, [], [f803, f2188, f2184])).
fof(f803, plain, (sP49 | ~ sP126), inference(cnf_transformation, [], [f381])).
fof(f2182, plain, (~ spl130_183 | spl130_208), inference(avatar_split_clause, [], [f754, f2179, f2055])).
fof(f754, plain, (sP50 | ~ sP127), inference(cnf_transformation, [], [f380])).
fof(f380, plain, ((sP74 & sP73 & sP72 & sP71 & sP70 & sP69 & sP68 & sP67 & sP66 & sP65 & sP64 & sP63 & sP62 & sP61 & sP60 & sP59 & sP58 & sP57 & sP56 & sP55 & sP54 & sP53 & sP52 & sP51 & sP50) | ~ sP127), inference(nnf_transformation, [], [f374])).
fof(f2177, plain, (~ spl130_183 | spl130_207), inference(avatar_split_clause, [], [f755, f2174, f2055])).
fof(f755, plain, (sP51 | ~ sP127), inference(cnf_transformation, [], [f380])).
fof(f2167, plain, (~ spl130_183 | spl130_205), inference(avatar_split_clause, [], [f757, f2164, f2055])).
fof(f757, plain, (sP53 | ~ sP127), inference(cnf_transformation, [], [f380])).
fof(f2162, plain, (~ spl130_183 | spl130_204), inference(avatar_split_clause, [], [f758, f2159, f2055])).
fof(f758, plain, (sP54 | ~ sP127), inference(cnf_transformation, [], [f380])).
fof(f2157, plain, (~ spl130_183 | spl130_203), inference(avatar_split_clause, [], [f759, f2154, f2055])).
fof(f759, plain, (sP55 | ~ sP127), inference(cnf_transformation, [], [f380])).
fof(f2152, plain, (~ spl130_183 | spl130_202), inference(avatar_split_clause, [], [f760, f2149, f2055])).
fof(f760, plain, (sP56 | ~ sP127), inference(cnf_transformation, [], [f380])).
fof(f2142, plain, (~ spl130_183 | spl130_200), inference(avatar_split_clause, [], [f762, f2139, f2055])).
fof(f762, plain, (sP58 | ~ sP127), inference(cnf_transformation, [], [f380])).
fof(f2137, plain, (~ spl130_183 | spl130_199), inference(avatar_split_clause, [], [f763, f2134, f2055])).
fof(f763, plain, (sP59 | ~ sP127), inference(cnf_transformation, [], [f380])).
fof(f2132, plain, (~ spl130_183 | spl130_198), inference(avatar_split_clause, [], [f764, f2129, f2055])).
fof(f764, plain, (sP60 | ~ sP127), inference(cnf_transformation, [], [f380])).
fof(f2127, plain, (~ spl130_183 | spl130_197), inference(avatar_split_clause, [], [f765, f2124, f2055])).
fof(f765, plain, (sP61 | ~ sP127), inference(cnf_transformation, [], [f380])).
fof(f2122, plain, (~ spl130_183 | spl130_196), inference(avatar_split_clause, [], [f766, f2119, f2055])).
fof(f766, plain, (sP62 | ~ sP127), inference(cnf_transformation, [], [f380])).
fof(f2117, plain, (~ spl130_183 | spl130_195), inference(avatar_split_clause, [], [f767, f2114, f2055])).
fof(f767, plain, (sP63 | ~ sP127), inference(cnf_transformation, [], [f380])).
fof(f2112, plain, (~ spl130_183 | spl130_194), inference(avatar_split_clause, [], [f768, f2109, f2055])).
fof(f768, plain, (sP64 | ~ sP127), inference(cnf_transformation, [], [f380])).
fof(f2107, plain, (~ spl130_183 | spl130_193), inference(avatar_split_clause, [], [f769, f2104, f2055])).
fof(f769, plain, (sP65 | ~ sP127), inference(cnf_transformation, [], [f380])).
fof(f2092, plain, (~ spl130_183 | spl130_190), inference(avatar_split_clause, [], [f772, f2089, f2055])).
fof(f772, plain, (sP68 | ~ sP127), inference(cnf_transformation, [], [f380])).
fof(f2087, plain, (~ spl130_183 | spl130_189), inference(avatar_split_clause, [], [f773, f2084, f2055])).
fof(f773, plain, (sP69 | ~ sP127), inference(cnf_transformation, [], [f380])).
fof(f2082, plain, (~ spl130_183 | spl130_188), inference(avatar_split_clause, [], [f774, f2079, f2055])).
fof(f774, plain, (sP70 | ~ sP127), inference(cnf_transformation, [], [f380])).
fof(f2067, plain, (~ spl130_183 | spl130_185), inference(avatar_split_clause, [], [f777, f2064, f2055])).
fof(f777, plain, (sP73 | ~ sP127), inference(cnf_transformation, [], [f380])).
fof(f2062, plain, (~ spl130_183 | spl130_184), inference(avatar_split_clause, [], [f778, f2059, f2055])).
fof(f778, plain, (sP74 | ~ sP127), inference(cnf_transformation, [], [f380])).
fof(f2053, plain, (~ spl130_157 | spl130_182), inference(avatar_split_clause, [], [f729, f2050, f1926])).
fof(f729, plain, (sP75 | ~ sP128), inference(cnf_transformation, [], [f379])).
fof(f379, plain, ((sP99 & sP98 & sP97 & sP96 & sP95 & sP94 & sP93 & sP92 & sP91 & sP90 & sP89 & sP88 & sP87 & sP86 & sP85 & sP84 & sP83 & sP82 & sP81 & sP80 & sP79 & sP78 & sP77 & sP76 & sP75) | ~ sP128), inference(nnf_transformation, [], [f375])).
fof(f2048, plain, (~ spl130_157 | spl130_181), inference(avatar_split_clause, [], [f730, f2045, f1926])).
fof(f730, plain, (sP76 | ~ sP128), inference(cnf_transformation, [], [f379])).
fof(f2043, plain, (~ spl130_157 | spl130_180), inference(avatar_split_clause, [], [f731, f2040, f1926])).
fof(f731, plain, (sP77 | ~ sP128), inference(cnf_transformation, [], [f379])).
fof(f2033, plain, (~ spl130_157 | spl130_178), inference(avatar_split_clause, [], [f733, f2030, f1926])).
fof(f733, plain, (sP79 | ~ sP128), inference(cnf_transformation, [], [f379])).
fof(f2028, plain, (~ spl130_157 | spl130_177), inference(avatar_split_clause, [], [f734, f2025, f1926])).
fof(f734, plain, (sP80 | ~ sP128), inference(cnf_transformation, [], [f379])).
fof(f2023, plain, (~ spl130_157 | spl130_176), inference(avatar_split_clause, [], [f735, f2020, f1926])).
fof(f735, plain, (sP81 | ~ sP128), inference(cnf_transformation, [], [f379])).
fof(f2018, plain, (~ spl130_157 | spl130_175), inference(avatar_split_clause, [], [f736, f2015, f1926])).
fof(f736, plain, (sP82 | ~ sP128), inference(cnf_transformation, [], [f379])).
fof(f2008, plain, (~ spl130_157 | spl130_173), inference(avatar_split_clause, [], [f738, f2005, f1926])).
fof(f738, plain, (sP84 | ~ sP128), inference(cnf_transformation, [], [f379])).
fof(f2003, plain, (~ spl130_157 | spl130_172), inference(avatar_split_clause, [], [f739, f2000, f1926])).
fof(f739, plain, (sP85 | ~ sP128), inference(cnf_transformation, [], [f379])).
fof(f1998, plain, (~ spl130_157 | spl130_171), inference(avatar_split_clause, [], [f740, f1995, f1926])).
fof(f740, plain, (sP86 | ~ sP128), inference(cnf_transformation, [], [f379])).
fof(f1993, plain, (~ spl130_157 | spl130_170), inference(avatar_split_clause, [], [f741, f1990, f1926])).
fof(f741, plain, (sP87 | ~ sP128), inference(cnf_transformation, [], [f379])).
fof(f1983, plain, (~ spl130_157 | spl130_168), inference(avatar_split_clause, [], [f743, f1980, f1926])).
fof(f743, plain, (sP89 | ~ sP128), inference(cnf_transformation, [], [f379])).
fof(f1978, plain, (~ spl130_157 | spl130_167), inference(avatar_split_clause, [], [f744, f1975, f1926])).
fof(f744, plain, (sP90 | ~ sP128), inference(cnf_transformation, [], [f379])).
fof(f1973, plain, (~ spl130_157 | spl130_166), inference(avatar_split_clause, [], [f745, f1970, f1926])).
fof(f745, plain, (sP91 | ~ sP128), inference(cnf_transformation, [], [f379])).
fof(f1968, plain, (~ spl130_157 | spl130_165), inference(avatar_split_clause, [], [f746, f1965, f1926])).
fof(f746, plain, (sP92 | ~ sP128), inference(cnf_transformation, [], [f379])).
fof(f1963, plain, (~ spl130_157 | spl130_164), inference(avatar_split_clause, [], [f747, f1960, f1926])).
fof(f747, plain, (sP93 | ~ sP128), inference(cnf_transformation, [], [f379])).
fof(f1958, plain, (~ spl130_157 | spl130_163), inference(avatar_split_clause, [], [f748, f1955, f1926])).
fof(f748, plain, (sP94 | ~ sP128), inference(cnf_transformation, [], [f379])).
fof(f1953, plain, (~ spl130_157 | spl130_162), inference(avatar_split_clause, [], [f749, f1950, f1926])).
fof(f749, plain, (sP95 | ~ sP128), inference(cnf_transformation, [], [f379])).
fof(f1943, plain, (~ spl130_157 | spl130_160), inference(avatar_split_clause, [], [f751, f1940, f1926])).
fof(f751, plain, (sP97 | ~ sP128), inference(cnf_transformation, [], [f379])).
fof(f1933, plain, (~ spl130_157 | spl130_158), inference(avatar_split_clause, [], [f753, f1930, f1926])).
fof(f753, plain, (sP99 | ~ sP128), inference(cnf_transformation, [], [f379])).
fof(f1924, plain, (~ spl130_131 | spl130_156), inference(avatar_split_clause, [], [f704, f1921, f1797])).
fof(f704, plain, (sP100 | ~ sP129), inference(cnf_transformation, [], [f378])).
fof(f378, plain, ((sP124 & sP123 & sP122 & sP121 & sP120 & sP119 & sP118 & sP117 & sP116 & sP115 & sP114 & sP113 & sP112 & sP111 & sP110 & sP109 & sP108 & sP107 & sP106 & sP105 & sP104 & sP103 & sP102 & sP101 & sP100) | ~ sP129), inference(nnf_transformation, [], [f376])).
fof(f1919, plain, (~ spl130_131 | spl130_155), inference(avatar_split_clause, [], [f705, f1916, f1797])).
fof(f705, plain, (sP101 | ~ sP129), inference(cnf_transformation, [], [f378])).
fof(f1914, plain, (~ spl130_131 | spl130_154), inference(avatar_split_clause, [], [f706, f1911, f1797])).
fof(f706, plain, (sP102 | ~ sP129), inference(cnf_transformation, [], [f378])).
fof(f1909, plain, (~ spl130_131 | spl130_153), inference(avatar_split_clause, [], [f707, f1906, f1797])).
fof(f707, plain, (sP103 | ~ sP129), inference(cnf_transformation, [], [f378])).
fof(f1899, plain, (~ spl130_131 | spl130_151), inference(avatar_split_clause, [], [f709, f1896, f1797])).
fof(f709, plain, (sP105 | ~ sP129), inference(cnf_transformation, [], [f378])).
fof(f1894, plain, (~ spl130_131 | spl130_150), inference(avatar_split_clause, [], [f710, f1891, f1797])).
fof(f710, plain, (sP106 | ~ sP129), inference(cnf_transformation, [], [f378])).
fof(f1889, plain, (~ spl130_131 | spl130_149), inference(avatar_split_clause, [], [f711, f1886, f1797])).
fof(f711, plain, (sP107 | ~ sP129), inference(cnf_transformation, [], [f378])).
fof(f1884, plain, (~ spl130_131 | spl130_148), inference(avatar_split_clause, [], [f712, f1881, f1797])).
fof(f712, plain, (sP108 | ~ sP129), inference(cnf_transformation, [], [f378])).
fof(f1874, plain, (~ spl130_131 | spl130_146), inference(avatar_split_clause, [], [f714, f1871, f1797])).
fof(f714, plain, (sP110 | ~ sP129), inference(cnf_transformation, [], [f378])).
fof(f1869, plain, (~ spl130_131 | spl130_145), inference(avatar_split_clause, [], [f715, f1866, f1797])).
fof(f715, plain, (sP111 | ~ sP129), inference(cnf_transformation, [], [f378])).
fof(f1864, plain, (~ spl130_131 | spl130_144), inference(avatar_split_clause, [], [f716, f1861, f1797])).
fof(f716, plain, (sP112 | ~ sP129), inference(cnf_transformation, [], [f378])).
fof(f1859, plain, (~ spl130_131 | spl130_143), inference(avatar_split_clause, [], [f717, f1856, f1797])).
fof(f717, plain, (sP113 | ~ sP129), inference(cnf_transformation, [], [f378])).
fof(f1844, plain, (~ spl130_131 | spl130_140), inference(avatar_split_clause, [], [f720, f1841, f1797])).
fof(f720, plain, (sP116 | ~ sP129), inference(cnf_transformation, [], [f378])).
fof(f1839, plain, (~ spl130_131 | spl130_139), inference(avatar_split_clause, [], [f721, f1836, f1797])).
fof(f721, plain, (sP117 | ~ sP129), inference(cnf_transformation, [], [f378])).
fof(f1834, plain, (~ spl130_131 | spl130_138), inference(avatar_split_clause, [], [f722, f1831, f1797])).
fof(f722, plain, (sP118 | ~ sP129), inference(cnf_transformation, [], [f378])).
fof(f1824, plain, (~ spl130_131 | spl130_136), inference(avatar_split_clause, [], [f724, f1821, f1797])).
fof(f724, plain, (sP120 | ~ sP129), inference(cnf_transformation, [], [f378])).
fof(f1819, plain, (~ spl130_131 | spl130_135), inference(avatar_split_clause, [], [f725, f1816, f1797])).
fof(f725, plain, (sP121 | ~ sP129), inference(cnf_transformation, [], [f378])).
fof(f1814, plain, (~ spl130_131 | spl130_134), inference(avatar_split_clause, [], [f726, f1811, f1797])).
fof(f726, plain, (sP122 | ~ sP129), inference(cnf_transformation, [], [f378])).
fof(f1809, plain, (~ spl130_131 | spl130_133), inference(avatar_split_clause, [], [f727, f1806, f1797])).
fof(f727, plain, (sP123 | ~ sP129), inference(cnf_transformation, [], [f378])).
fof(f1804, plain, (~ spl130_131 | spl130_132), inference(avatar_split_clause, [], [f728, f1801, f1797])).
fof(f728, plain, (sP124 | ~ sP129), inference(cnf_transformation, [], [f378])).
fof(f1793, plain, (spl130_122 | spl130_117 | spl130_112 | spl130_107 | spl130_102), inference(avatar_split_clause, [], [f546, f1625, f1646, f1667, f1688, f1709])).
fof(f546, plain, ((e1 = op(e0, e4)) | (e1 = op(e0, e3)) | (e1 = op(e0, e2)) | (e1 = op(e0, e1)) | (op(e0, e0) = e1)), inference(cnf_transformation, [], [f3])).
fof(f3, plain, (((e4 = op(e4, e4)) | (e4 = op(e3, e4)) | (e4 = op(e2, e4)) | (e4 = op(e1, e4)) | (e4 = op(e0, e4))) & ((e4 = op(e4, e4)) | (e4 = op(e4, e3)) | (e4 = op(e4, e2)) | (e4 = op(e4, e1)) | (e4 = op(e4, e0))) & ((e3 = op(e4, e4)) | (e3 = op(e3, e4)) | (e3 = op(e2, e4)) | (e3 = op(e1, e4)) | (e3 = op(e0, e4))) & ((e3 = op(e4, e4)) | (e3 = op(e4, e3)) | (e3 = op(e4, e2)) | (e3 = op(e4, e1)) | (e3 = op(e4, e0))) & ((e2 = op(e4, e4)) | (e2 = op(e3, e4)) | (e2 = op(e2, e4)) | (e2 = op(e1, e4)) | (e2 = op(e0, e4))) & ((e2 = op(e4, e4)) | (e2 = op(e4, e3)) | (e2 = op(e4, e2)) | (e2 = op(e4, e1)) | (e2 = op(e4, e0))) & ((e1 = op(e4, e4)) | (e1 = op(e3, e4)) | (e1 = op(e2, e4)) | (e1 = op(e1, e4)) | (e1 = op(e0, e4))) & ((e1 = op(e4, e4)) | (e1 = op(e4, e3)) | (e1 = op(e4, e2)) | (e1 = op(e4, e1)) | (e1 = op(e4, e0))) & ((e0 = op(e4, e4)) | (e0 = op(e3, e4)) | (e0 = op(e2, e4)) | (e0 = op(e1, e4)) | (e0 = op(e0, e4))) & ((e0 = op(e4, e4)) | (e0 = op(e4, e3)) | (e0 = op(e4, e2)) | (e0 = op(e4, e1)) | (e0 = op(e4, e0))) & ((e4 = op(e4, e3)) | (e4 = op(e3, e3)) | (e4 = op(e2, e3)) | (e4 = op(e1, e3)) | (e4 = op(e0, e3))) & ((e4 = op(e3, e4)) | (e4 = op(e3, e3)) | (e4 = op(e3, e2)) | (e4 = op(e3, e1)) | (e4 = op(e3, e0))) & ((e3 = op(e4, e3)) | (e3 = op(e3, e3)) | (e3 = op(e2, e3)) | (e3 = op(e1, e3)) | (e3 = op(e0, e3))) & ((e3 = op(e3, e4)) | (e3 = op(e3, e3)) | (e3 = op(e3, e2)) | (e3 = op(e3, e1)) | (e3 = op(e3, e0))) & ((e2 = op(e4, e3)) | (e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(e0, e3))) & ((e2 = op(e3, e4)) | (e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, e0))) & ((e1 = op(e4, e3)) | (e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(e0, e3))) & ((e1 = op(e3, e4)) | (e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, e0))) & ((e0 = op(e4, e3)) | (e0 = op(e3, e3)) | (e0 = op(e2, e3)) | (e0 = op(e1, e3)) | (e0 = op(e0, e3))) & ((e0 = op(e3, e4)) | (e0 = op(e3, e3)) | (e0 = op(e3, e2)) | (e0 = op(e3, e1)) | (e0 = op(e3, e0))) & ((e4 = op(e4, e2)) | (e4 = op(e3, e2)) | (e4 = op(e2, e2)) | (e4 = op(e1, e2)) | (e4 = op(e0, e2))) & ((e4 = op(e2, e4)) | (e4 = op(e2, e3)) | (e4 = op(e2, e2)) | (e4 = op(e2, e1)) | (e4 = op(e2, e0))) & ((e3 = op(e4, e2)) | (e3 = op(e3, e2)) | (e3 = op(e2, e2)) | (e3 = op(e1, e2)) | (e3 = op(e0, e2))) & ((e3 = op(e2, e4)) | (e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, e0))) & ((e2 = op(e4, e2)) | (e2 = op(e3, e2)) | (e2 = op(e2, e2)) | (e2 = op(e1, e2)) | (e2 = op(e0, e2))) & ((e2 = op(e2, e4)) | (e2 = op(e2, e3)) | (e2 = op(e2, e2)) | (e2 = op(e2, e1)) | (e2 = op(e2, e0))) & ((e1 = op(e4, e2)) | (e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(e0, e2))) & ((e1 = op(e2, e4)) | (e1 = op(e2, e3)) | (e1 = op(e2, e2)) | (e1 = op(e2, e1)) | (e1 = op(e2, e0))) & ((e0 = op(e4, e2)) | (e0 = op(e3, e2)) | (e0 = op(e2, e2)) | (e0 = op(e1, e2)) | (e0 = op(e0, e2))) & ((e0 = op(e2, e4)) | (e0 = op(e2, e3)) | (e0 = op(e2, e2)) | (e0 = op(e2, e1)) | (e0 = op(e2, e0))) & ((e4 = op(e4, e1)) | (e4 = op(e3, e1)) | (e4 = op(e2, e1)) | (e4 = op(e1, e1)) | (e4 = op(e0, e1))) & ((e4 = op(e1, e4)) | (e4 = op(e1, e3)) | (e4 = op(e1, e2)) | (e4 = op(e1, e1)) | (e4 = op(e1, e0))) & ((e3 = op(e4, e1)) | (e3 = op(e3, e1)) | (e3 = op(e2, e1)) | (e3 = op(e1, e1)) | (e3 = op(e0, e1))) & ((e3 = op(e1, e4)) | (e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, e0))) & ((e2 = op(e4, e1)) | (e2 = op(e3, e1)) | (e2 = op(e2, e1)) | (e2 = op(e1, e1)) | (e2 = op(e0, e1))) & ((e2 = op(e1, e4)) | (e2 = op(e1, e3)) | (e2 = op(e1, e2)) | (e2 = op(e1, e1)) | (e2 = op(e1, e0))) & ((e1 = op(e4, e1)) | (e1 = op(e3, e1)) | (e1 = op(e2, e1)) | (e1 = op(e1, e1)) | (e1 = op(e0, e1))) & ((e1 = op(e1, e4)) | (e1 = op(e1, e3)) | (e1 = op(e1, e2)) | (e1 = op(e1, e1)) | (e1 = op(e1, e0))) & ((e0 = op(e4, e1)) | (e0 = op(e3, e1)) | (e0 = op(e2, e1)) | (e0 = op(e1, e1)) | (e0 = op(e0, e1))) & ((e0 = op(e1, e4)) | (e0 = op(e1, e3)) | (e0 = op(e1, e2)) | (e0 = op(e1, e1)) | (e0 = op(e1, e0))) & ((e4 = op(e4, e0)) | (e4 = op(e3, e0)) | (e4 = op(e2, e0)) | (e4 = op(e1, e0)) | (op(e0, e0) = e4)) & ((e4 = op(e0, e4)) | (e4 = op(e0, e3)) | (e4 = op(e0, e2)) | (e4 = op(e0, e1)) | (op(e0, e0) = e4)) & ((e3 = op(e4, e0)) | (e3 = op(e3, e0)) | (e3 = op(e2, e0)) | (e3 = op(e1, e0)) | (op(e0, e0) = e3)) & ((e3 = op(e0, e4)) | (e3 = op(e0, e3)) | (e3 = op(e0, e2)) | (e3 = op(e0, e1)) | (op(e0, e0) = e3)) & ((e2 = op(e4, e0)) | (e2 = op(e3, e0)) | (e2 = op(e2, e0)) | (e2 = op(e1, e0)) | (op(e0, e0) = e2)) & ((e2 = op(e0, e4)) | (e2 = op(e0, e3)) | (e2 = op(e0, e2)) | (e2 = op(e0, e1)) | (op(e0, e0) = e2)) & ((e1 = op(e4, e0)) | (e1 = op(e3, e0)) | (e1 = op(e2, e0)) | (e1 = op(e1, e0)) | (op(e0, e0) = e1)) & ((e1 = op(e0, e4)) | (e1 = op(e0, e3)) | (e1 = op(e0, e2)) | (e1 = op(e0, e1)) | (op(e0, e0) = e1)) & ((e0 = op(e4, e0)) | (e0 = op(e3, e0)) | (e0 = op(e2, e0)) | (e0 = op(e1, e0)) | (e0 = op(e0, e0))) & ((e0 = op(e0, e4)) | (e0 = op(e0, e3)) | (e0 = op(e0, e2)) | (e0 = op(e0, e1)) | (e0 = op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG099+1.p', ax3)).
fof(f1792, plain, (spl130_122 | spl130_97 | spl130_72 | spl130_47 | spl130_22), inference(avatar_split_clause, [], [f547, f1289, f1394, f1499, f1604, f1709])).
fof(f547, plain, ((e1 = op(e4, e0)) | (e1 = op(e3, e0)) | (e1 = op(e2, e0)) | (e1 = op(e1, e0)) | (op(e0, e0) = e1)), inference(cnf_transformation, [], [f3])).
fof(f1791, plain, (spl130_123 | spl130_118 | spl130_113 | spl130_108 | spl130_103), inference(avatar_split_clause, [], [f548, f1629, f1650, f1671, f1692, f1713])).
fof(f548, plain, ((e2 = op(e0, e4)) | (e2 = op(e0, e3)) | (e2 = op(e0, e2)) | (e2 = op(e0, e1)) | (op(e0, e0) = e2)), inference(cnf_transformation, [], [f3])).
fof(f1790, plain, (spl130_123 | spl130_98 | spl130_73 | spl130_48 | spl130_23), inference(avatar_split_clause, [], [f549, f1293, f1398, f1503, f1608, f1713])).
fof(f549, plain, ((e2 = op(e4, e0)) | (e2 = op(e3, e0)) | (e2 = op(e2, e0)) | (e2 = op(e1, e0)) | (op(e0, e0) = e2)), inference(cnf_transformation, [], [f3])).
fof(f1789, plain, (spl130_124 | spl130_119 | spl130_114 | spl130_109 | spl130_104), inference(avatar_split_clause, [], [f550, f1633, f1654, f1675, f1696, f1717])).
fof(f550, plain, ((e3 = op(e0, e4)) | (e3 = op(e0, e3)) | (e3 = op(e0, e2)) | (e3 = op(e0, e1)) | (op(e0, e0) = e3)), inference(cnf_transformation, [], [f3])).
fof(f1788, plain, (spl130_124 | spl130_99 | spl130_74 | spl130_49 | spl130_24), inference(avatar_split_clause, [], [f551, f1297, f1402, f1507, f1612, f1717])).
fof(f551, plain, ((e3 = op(e4, e0)) | (e3 = op(e3, e0)) | (e3 = op(e2, e0)) | (e3 = op(e1, e0)) | (op(e0, e0) = e3)), inference(cnf_transformation, [], [f3])).
fof(f1787, plain, (spl130_125 | spl130_120 | spl130_115 | spl130_110 | spl130_105), inference(avatar_split_clause, [], [f552, f1637, f1658, f1679, f1700, f1721])).
fof(f552, plain, ((e4 = op(e0, e4)) | (e4 = op(e0, e3)) | (e4 = op(e0, e2)) | (e4 = op(e0, e1)) | (op(e0, e0) = e4)), inference(cnf_transformation, [], [f3])).
fof(f1786, plain, (spl130_125 | spl130_100 | spl130_75 | spl130_50 | spl130_25), inference(avatar_split_clause, [], [f553, f1301, f1406, f1511, f1616, f1721])).
fof(f553, plain, ((e4 = op(e4, e0)) | (e4 = op(e3, e0)) | (e4 = op(e2, e0)) | (e4 = op(e1, e0)) | (op(e0, e0) = e4)), inference(cnf_transformation, [], [f3])).
fof(f1785, plain, (spl130_96 | spl130_91 | spl130_86 | spl130_81 | spl130_76), inference(avatar_split_clause, [], [f554, f1516, f1537, f1558, f1579, f1600])).
fof(f554, plain, ((e0 = op(e1, e4)) | (e0 = op(e1, e3)) | (e0 = op(e1, e2)) | (e0 = op(e1, e1)) | (e0 = op(e1, e0))), inference(cnf_transformation, [], [f3])).
fof(f1784, plain, (spl130_116 | spl130_91 | spl130_66 | spl130_41 | spl130_16), inference(avatar_split_clause, [], [f555, f1264, f1369, f1474, f1579, f1684])).
fof(f555, plain, ((e0 = op(e4, e1)) | (e0 = op(e3, e1)) | (e0 = op(e2, e1)) | (e0 = op(e1, e1)) | (e0 = op(e0, e1))), inference(cnf_transformation, [], [f3])).
fof(f1781, plain, (spl130_98 | spl130_93 | spl130_88 | spl130_83 | spl130_78), inference(avatar_split_clause, [], [f558, f1524, f1545, f1566, f1587, f1608])).
fof(f558, plain, ((e2 = op(e1, e4)) | (e2 = op(e1, e3)) | (e2 = op(e1, e2)) | (e2 = op(e1, e1)) | (e2 = op(e1, e0))), inference(cnf_transformation, [], [f3])).
fof(f1780, plain, (spl130_118 | spl130_93 | spl130_68 | spl130_43 | spl130_18), inference(avatar_split_clause, [], [f559, f1272, f1377, f1482, f1587, f1692])).
fof(f559, plain, ((e2 = op(e4, e1)) | (e2 = op(e3, e1)) | (e2 = op(e2, e1)) | (e2 = op(e1, e1)) | (e2 = op(e0, e1))), inference(cnf_transformation, [], [f3])).
fof(f1779, plain, (spl130_99 | spl130_94 | spl130_89 | spl130_84 | spl130_79), inference(avatar_split_clause, [], [f560, f1528, f1549, f1570, f1591, f1612])).
fof(f560, plain, ((e3 = op(e1, e4)) | (e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, e0))), inference(cnf_transformation, [], [f3])).
fof(f1778, plain, (spl130_119 | spl130_94 | spl130_69 | spl130_44 | spl130_19), inference(avatar_split_clause, [], [f561, f1276, f1381, f1486, f1591, f1696])).
fof(f561, plain, ((e3 = op(e4, e1)) | (e3 = op(e3, e1)) | (e3 = op(e2, e1)) | (e3 = op(e1, e1)) | (e3 = op(e0, e1))), inference(cnf_transformation, [], [f3])).
fof(f1777, plain, (spl130_100 | spl130_95 | spl130_90 | spl130_85 | spl130_80), inference(avatar_split_clause, [], [f562, f1532, f1553, f1574, f1595, f1616])).
fof(f562, plain, ((e4 = op(e1, e4)) | (e4 = op(e1, e3)) | (e4 = op(e1, e2)) | (e4 = op(e1, e1)) | (e4 = op(e1, e0))), inference(cnf_transformation, [], [f3])).
fof(f1776, plain, (spl130_120 | spl130_95 | spl130_70 | spl130_45 | spl130_20), inference(avatar_split_clause, [], [f563, f1280, f1385, f1490, f1595, f1700])).
fof(f563, plain, ((e4 = op(e4, e1)) | (e4 = op(e3, e1)) | (e4 = op(e2, e1)) | (e4 = op(e1, e1)) | (e4 = op(e0, e1))), inference(cnf_transformation, [], [f3])).
fof(f1775, plain, (spl130_71 | spl130_66 | spl130_61 | spl130_56 | spl130_51), inference(avatar_split_clause, [], [f564, f1411, f1432, f1453, f1474, f1495])).
fof(f564, plain, ((e0 = op(e2, e4)) | (e0 = op(e2, e3)) | (e0 = op(e2, e2)) | (e0 = op(e2, e1)) | (e0 = op(e2, e0))), inference(cnf_transformation, [], [f3])).
fof(f1774, plain, (spl130_111 | spl130_86 | spl130_61 | spl130_36 | spl130_11), inference(avatar_split_clause, [], [f565, f1243, f1348, f1453, f1558, f1663])).
fof(f565, plain, ((e0 = op(e4, e2)) | (e0 = op(e3, e2)) | (e0 = op(e2, e2)) | (e0 = op(e1, e2)) | (e0 = op(e0, e2))), inference(cnf_transformation, [], [f3])).
fof(f1773, plain, (spl130_72 | spl130_67 | spl130_62 | spl130_57 | spl130_52), inference(avatar_split_clause, [], [f566, f1415, f1436, f1457, f1478, f1499])).
fof(f566, plain, ((e1 = op(e2, e4)) | (e1 = op(e2, e3)) | (e1 = op(e2, e2)) | (e1 = op(e2, e1)) | (e1 = op(e2, e0))), inference(cnf_transformation, [], [f3])).
fof(f1772, plain, (spl130_112 | spl130_87 | spl130_62 | spl130_37 | spl130_12), inference(avatar_split_clause, [], [f567, f1247, f1352, f1457, f1562, f1667])).
fof(f567, plain, ((e1 = op(e4, e2)) | (e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(e0, e2))), inference(cnf_transformation, [], [f3])).
fof(f1769, plain, (spl130_74 | spl130_69 | spl130_64 | spl130_59 | spl130_54), inference(avatar_split_clause, [], [f570, f1423, f1444, f1465, f1486, f1507])).
fof(f570, plain, ((e3 = op(e2, e4)) | (e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, e0))), inference(cnf_transformation, [], [f3])).
fof(f1768, plain, (spl130_114 | spl130_89 | spl130_64 | spl130_39 | spl130_14), inference(avatar_split_clause, [], [f571, f1255, f1360, f1465, f1570, f1675])).
fof(f571, plain, ((e3 = op(e4, e2)) | (e3 = op(e3, e2)) | (e3 = op(e2, e2)) | (e3 = op(e1, e2)) | (e3 = op(e0, e2))), inference(cnf_transformation, [], [f3])).
fof(f1767, plain, (spl130_75 | spl130_70 | spl130_65 | spl130_60 | spl130_55), inference(avatar_split_clause, [], [f572, f1427, f1448, f1469, f1490, f1511])).
fof(f572, plain, ((e4 = op(e2, e4)) | (e4 = op(e2, e3)) | (e4 = op(e2, e2)) | (e4 = op(e2, e1)) | (e4 = op(e2, e0))), inference(cnf_transformation, [], [f3])).
fof(f1766, plain, (spl130_115 | spl130_90 | spl130_65 | spl130_40 | spl130_15), inference(avatar_split_clause, [], [f573, f1259, f1364, f1469, f1574, f1679])).
fof(f573, plain, ((e4 = op(e4, e2)) | (e4 = op(e3, e2)) | (e4 = op(e2, e2)) | (e4 = op(e1, e2)) | (e4 = op(e0, e2))), inference(cnf_transformation, [], [f3])).
fof(f1765, plain, (spl130_46 | spl130_41 | spl130_36 | spl130_31 | spl130_26), inference(avatar_split_clause, [], [f574, f1306, f1327, f1348, f1369, f1390])).
fof(f574, plain, ((e0 = op(e3, e4)) | (e0 = op(e3, e3)) | (e0 = op(e3, e2)) | (e0 = op(e3, e1)) | (e0 = op(e3, e0))), inference(cnf_transformation, [], [f3])).
fof(f1764, plain, (spl130_106 | spl130_81 | spl130_56 | spl130_31 | spl130_6), inference(avatar_split_clause, [], [f575, f1222, f1327, f1432, f1537, f1642])).
fof(f575, plain, ((e0 = op(e4, e3)) | (e0 = op(e3, e3)) | (e0 = op(e2, e3)) | (e0 = op(e1, e3)) | (e0 = op(e0, e3))), inference(cnf_transformation, [], [f3])).
fof(f1763, plain, (spl130_47 | spl130_42 | spl130_37 | spl130_32 | spl130_27), inference(avatar_split_clause, [], [f576, f1310, f1331, f1352, f1373, f1394])).
fof(f576, plain, ((e1 = op(e3, e4)) | (e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, e0))), inference(cnf_transformation, [], [f3])).
fof(f1762, plain, (spl130_107 | spl130_82 | spl130_57 | spl130_32 | spl130_7), inference(avatar_split_clause, [], [f577, f1226, f1331, f1436, f1541, f1646])).
fof(f577, plain, ((e1 = op(e4, e3)) | (e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(e0, e3))), inference(cnf_transformation, [], [f3])).
fof(f1761, plain, (spl130_48 | spl130_43 | spl130_38 | spl130_33 | spl130_28), inference(avatar_split_clause, [], [f578, f1314, f1335, f1356, f1377, f1398])).
fof(f578, plain, ((e2 = op(e3, e4)) | (e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, e0))), inference(cnf_transformation, [], [f3])).
fof(f1760, plain, (spl130_108 | spl130_83 | spl130_58 | spl130_33 | spl130_8), inference(avatar_split_clause, [], [f579, f1230, f1335, f1440, f1545, f1650])).
fof(f579, plain, ((e2 = op(e4, e3)) | (e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(e0, e3))), inference(cnf_transformation, [], [f3])).
fof(f1757, plain, (spl130_50 | spl130_45 | spl130_40 | spl130_35 | spl130_30), inference(avatar_split_clause, [], [f582, f1322, f1343, f1364, f1385, f1406])).
fof(f582, plain, ((e4 = op(e3, e4)) | (e4 = op(e3, e3)) | (e4 = op(e3, e2)) | (e4 = op(e3, e1)) | (e4 = op(e3, e0))), inference(cnf_transformation, [], [f3])).
fof(f1756, plain, (spl130_110 | spl130_85 | spl130_60 | spl130_35 | spl130_10), inference(avatar_split_clause, [], [f583, f1238, f1343, f1448, f1553, f1658])).
fof(f583, plain, ((e4 = op(e4, e3)) | (e4 = op(e3, e3)) | (e4 = op(e2, e3)) | (e4 = op(e1, e3)) | (e4 = op(e0, e3))), inference(cnf_transformation, [], [f3])).
fof(f1755, plain, (spl130_21 | spl130_16 | spl130_11 | spl130_6 | spl130_1), inference(avatar_split_clause, [], [f584, f1201, f1222, f1243, f1264, f1285])).
fof(f584, plain, ((e0 = op(e4, e4)) | (e0 = op(e4, e3)) | (e0 = op(e4, e2)) | (e0 = op(e4, e1)) | (e0 = op(e4, e0))), inference(cnf_transformation, [], [f3])).
fof(f1754, plain, (spl130_101 | spl130_76 | spl130_51 | spl130_26 | spl130_1), inference(avatar_split_clause, [], [f585, f1201, f1306, f1411, f1516, f1621])).
fof(f585, plain, ((e0 = op(e4, e4)) | (e0 = op(e3, e4)) | (e0 = op(e2, e4)) | (e0 = op(e1, e4)) | (e0 = op(e0, e4))), inference(cnf_transformation, [], [f3])).
fof(f1753, plain, (spl130_22 | spl130_17 | spl130_12 | spl130_7 | spl130_2), inference(avatar_split_clause, [], [f586, f1205, f1226, f1247, f1268, f1289])).
fof(f586, plain, ((e1 = op(e4, e4)) | (e1 = op(e4, e3)) | (e1 = op(e4, e2)) | (e1 = op(e4, e1)) | (e1 = op(e4, e0))), inference(cnf_transformation, [], [f3])).
fof(f1752, plain, (spl130_102 | spl130_77 | spl130_52 | spl130_27 | spl130_2), inference(avatar_split_clause, [], [f587, f1205, f1310, f1415, f1520, f1625])).
fof(f587, plain, ((e1 = op(e4, e4)) | (e1 = op(e3, e4)) | (e1 = op(e2, e4)) | (e1 = op(e1, e4)) | (e1 = op(e0, e4))), inference(cnf_transformation, [], [f3])).
fof(f1751, plain, (spl130_23 | spl130_18 | spl130_13 | spl130_8 | spl130_3), inference(avatar_split_clause, [], [f588, f1209, f1230, f1251, f1272, f1293])).
fof(f588, plain, ((e2 = op(e4, e4)) | (e2 = op(e4, e3)) | (e2 = op(e4, e2)) | (e2 = op(e4, e1)) | (e2 = op(e4, e0))), inference(cnf_transformation, [], [f3])).
fof(f1750, plain, (spl130_103 | spl130_78 | spl130_53 | spl130_28 | spl130_3), inference(avatar_split_clause, [], [f589, f1209, f1314, f1419, f1524, f1629])).
fof(f589, plain, ((e2 = op(e4, e4)) | (e2 = op(e3, e4)) | (e2 = op(e2, e4)) | (e2 = op(e1, e4)) | (e2 = op(e0, e4))), inference(cnf_transformation, [], [f3])).
fof(f1749, plain, (spl130_24 | spl130_19 | spl130_14 | spl130_9 | spl130_4), inference(avatar_split_clause, [], [f590, f1213, f1234, f1255, f1276, f1297])).
fof(f590, plain, ((e3 = op(e4, e4)) | (e3 = op(e4, e3)) | (e3 = op(e4, e2)) | (e3 = op(e4, e1)) | (e3 = op(e4, e0))), inference(cnf_transformation, [], [f3])).
fof(f1748, plain, (spl130_104 | spl130_79 | spl130_54 | spl130_29 | spl130_4), inference(avatar_split_clause, [], [f591, f1213, f1318, f1423, f1528, f1633])).
fof(f591, plain, ((e3 = op(e4, e4)) | (e3 = op(e3, e4)) | (e3 = op(e2, e4)) | (e3 = op(e1, e4)) | (e3 = op(e0, e4))), inference(cnf_transformation, [], [f3])).
fof(f1745, plain, (spl130_126 | spl130_127 | spl130_128 | spl130_129 | spl130_130), inference(avatar_split_clause, [], [f543, f1742, f1738, f1734, f1730, f1726])).
fof(f543, plain, ((e4 = unit) | (e3 = unit) | (e2 = unit) | (e1 = unit) | (e0 = unit)), inference(cnf_transformation, [], [f2])).
fof(f1724, plain, (spl130_121 | spl130_122 | spl130_123 | spl130_124 | spl130_125), inference(avatar_split_clause, [], [f508, f1721, f1717, f1713, f1709, f1705])).
fof(f508, plain, ((op(e0, e0) = e4) | (op(e0, e0) = e3) | (op(e0, e0) = e2) | (op(e0, e0) = e1) | (e0 = op(e0, e0))), inference(cnf_transformation, [], [f1])).
fof(f1, plain, (((e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | (e1 = op(e4, e4)) | (e0 = op(e4, e4))) & ((e4 = op(e4, e3)) | (e3 = op(e4, e3)) | (e2 = op(e4, e3)) | (e1 = op(e4, e3)) | (e0 = op(e4, e3))) & ((e4 = op(e4, e2)) | (e3 = op(e4, e2)) | (e2 = op(e4, e2)) | (e1 = op(e4, e2)) | (e0 = op(e4, e2))) & ((e4 = op(e4, e1)) | (e3 = op(e4, e1)) | (e2 = op(e4, e1)) | (e1 = op(e4, e1)) | (e0 = op(e4, e1))) & ((e4 = op(e4, e0)) | (e3 = op(e4, e0)) | (e2 = op(e4, e0)) | (e1 = op(e4, e0)) | (e0 = op(e4, e0))) & ((e4 = op(e3, e4)) | (e3 = op(e3, e4)) | (e2 = op(e3, e4)) | (e1 = op(e3, e4)) | (e0 = op(e3, e4))) & ((e4 = op(e3, e3)) | (e3 = op(e3, e3)) | (e2 = op(e3, e3)) | (e1 = op(e3, e3)) | (e0 = op(e3, e3))) & ((e4 = op(e3, e2)) | (e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (e0 = op(e3, e2))) & ((e4 = op(e3, e1)) | (e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (e0 = op(e3, e1))) & ((e4 = op(e3, e0)) | (e3 = op(e3, e0)) | (e2 = op(e3, e0)) | (e1 = op(e3, e0)) | (e0 = op(e3, e0))) & ((e4 = op(e2, e4)) | (e3 = op(e2, e4)) | (e2 = op(e2, e4)) | (e1 = op(e2, e4)) | (e0 = op(e2, e4))) & ((e4 = op(e2, e3)) | (e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))) & ((e4 = op(e2, e2)) | (e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (e0 = op(e2, e2))) & ((e4 = op(e2, e1)) | (e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))) & ((e4 = op(e2, e0)) | (e3 = op(e2, e0)) | (e2 = op(e2, e0)) | (e1 = op(e2, e0)) | (e0 = op(e2, e0))) & ((e4 = op(e1, e4)) | (e3 = op(e1, e4)) | (e2 = op(e1, e4)) | (e1 = op(e1, e4)) | (e0 = op(e1, e4))) & ((e4 = op(e1, e3)) | (e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (e0 = op(e1, e3))) & ((e4 = op(e1, e2)) | (e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (e0 = op(e1, e2))) & ((e4 = op(e1, e1)) | (e3 = op(e1, e1)) | (e2 = op(e1, e1)) | (e1 = op(e1, e1)) | (e0 = op(e1, e1))) & ((e4 = op(e1, e0)) | (e3 = op(e1, e0)) | (e2 = op(e1, e0)) | (e1 = op(e1, e0)) | (e0 = op(e1, e0))) & ((e4 = op(e0, e4)) | (e3 = op(e0, e4)) | (e2 = op(e0, e4)) | (e1 = op(e0, e4)) | (e0 = op(e0, e4))) & ((e4 = op(e0, e3)) | (e3 = op(e0, e3)) | (e2 = op(e0, e3)) | (e1 = op(e0, e3)) | (e0 = op(e0, e3))) & ((e4 = op(e0, e2)) | (e3 = op(e0, e2)) | (e2 = op(e0, e2)) | (e1 = op(e0, e2)) | (e0 = op(e0, e2))) & ((e4 = op(e0, e1)) | (e3 = op(e0, e1)) | (e2 = op(e0, e1)) | (e1 = op(e0, e1)) | (e0 = op(e0, e1))) & ((op(e0, e0) = e4) | (op(e0, e0) = e3) | (op(e0, e0) = e2) | (op(e0, e0) = e1) | (e0 = op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG099+1.p', ax1)).
fof(f1703, plain, (spl130_116 | spl130_117 | spl130_118 | spl130_119 | spl130_120), inference(avatar_split_clause, [], [f509, f1700, f1696, f1692, f1688, f1684])).
fof(f509, plain, ((e4 = op(e0, e1)) | (e3 = op(e0, e1)) | (e2 = op(e0, e1)) | (e1 = op(e0, e1)) | (e0 = op(e0, e1))), inference(cnf_transformation, [], [f1])).
fof(f1682, plain, (spl130_111 | spl130_112 | spl130_113 | spl130_114 | spl130_115), inference(avatar_split_clause, [], [f510, f1679, f1675, f1671, f1667, f1663])).
fof(f510, plain, ((e4 = op(e0, e2)) | (e3 = op(e0, e2)) | (e2 = op(e0, e2)) | (e1 = op(e0, e2)) | (e0 = op(e0, e2))), inference(cnf_transformation, [], [f1])).
fof(f1661, plain, (spl130_106 | spl130_107 | spl130_108 | spl130_109 | spl130_110), inference(avatar_split_clause, [], [f511, f1658, f1654, f1650, f1646, f1642])).
fof(f511, plain, ((e4 = op(e0, e3)) | (e3 = op(e0, e3)) | (e2 = op(e0, e3)) | (e1 = op(e0, e3)) | (e0 = op(e0, e3))), inference(cnf_transformation, [], [f1])).
fof(f1640, plain, (spl130_101 | spl130_102 | spl130_103 | spl130_104 | spl130_105), inference(avatar_split_clause, [], [f512, f1637, f1633, f1629, f1625, f1621])).
fof(f512, plain, ((e4 = op(e0, e4)) | (e3 = op(e0, e4)) | (e2 = op(e0, e4)) | (e1 = op(e0, e4)) | (e0 = op(e0, e4))), inference(cnf_transformation, [], [f1])).
fof(f1619, plain, (spl130_96 | spl130_97 | spl130_98 | spl130_99 | spl130_100), inference(avatar_split_clause, [], [f513, f1616, f1612, f1608, f1604, f1600])).
fof(f513, plain, ((e4 = op(e1, e0)) | (e3 = op(e1, e0)) | (e2 = op(e1, e0)) | (e1 = op(e1, e0)) | (e0 = op(e1, e0))), inference(cnf_transformation, [], [f1])).
fof(f1598, plain, (spl130_91 | spl130_92 | spl130_93 | spl130_94 | spl130_95), inference(avatar_split_clause, [], [f514, f1595, f1591, f1587, f1583, f1579])).
fof(f514, plain, ((e4 = op(e1, e1)) | (e3 = op(e1, e1)) | (e2 = op(e1, e1)) | (e1 = op(e1, e1)) | (e0 = op(e1, e1))), inference(cnf_transformation, [], [f1])).
fof(f1577, plain, (spl130_86 | spl130_87 | spl130_88 | spl130_89 | spl130_90), inference(avatar_split_clause, [], [f515, f1574, f1570, f1566, f1562, f1558])).
fof(f515, plain, ((e4 = op(e1, e2)) | (e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (e0 = op(e1, e2))), inference(cnf_transformation, [], [f1])).
fof(f1556, plain, (spl130_81 | spl130_82 | spl130_83 | spl130_84 | spl130_85), inference(avatar_split_clause, [], [f516, f1553, f1549, f1545, f1541, f1537])).
fof(f516, plain, ((e4 = op(e1, e3)) | (e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (e0 = op(e1, e3))), inference(cnf_transformation, [], [f1])).
fof(f1535, plain, (spl130_76 | spl130_77 | spl130_78 | spl130_79 | spl130_80), inference(avatar_split_clause, [], [f517, f1532, f1528, f1524, f1520, f1516])).
fof(f517, plain, ((e4 = op(e1, e4)) | (e3 = op(e1, e4)) | (e2 = op(e1, e4)) | (e1 = op(e1, e4)) | (e0 = op(e1, e4))), inference(cnf_transformation, [], [f1])).
fof(f1514, plain, (spl130_71 | spl130_72 | spl130_73 | spl130_74 | spl130_75), inference(avatar_split_clause, [], [f518, f1511, f1507, f1503, f1499, f1495])).
fof(f518, plain, ((e4 = op(e2, e0)) | (e3 = op(e2, e0)) | (e2 = op(e2, e0)) | (e1 = op(e2, e0)) | (e0 = op(e2, e0))), inference(cnf_transformation, [], [f1])).
fof(f1493, plain, (spl130_66 | spl130_67 | spl130_68 | spl130_69 | spl130_70), inference(avatar_split_clause, [], [f519, f1490, f1486, f1482, f1478, f1474])).
fof(f519, plain, ((e4 = op(e2, e1)) | (e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))), inference(cnf_transformation, [], [f1])).
fof(f1472, plain, (spl130_61 | spl130_62 | spl130_63 | spl130_64 | spl130_65), inference(avatar_split_clause, [], [f520, f1469, f1465, f1461, f1457, f1453])).
fof(f520, plain, ((e4 = op(e2, e2)) | (e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (e0 = op(e2, e2))), inference(cnf_transformation, [], [f1])).
fof(f1451, plain, (spl130_56 | spl130_57 | spl130_58 | spl130_59 | spl130_60), inference(avatar_split_clause, [], [f521, f1448, f1444, f1440, f1436, f1432])).
fof(f521, plain, ((e4 = op(e2, e3)) | (e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))), inference(cnf_transformation, [], [f1])).
fof(f1430, plain, (spl130_51 | spl130_52 | spl130_53 | spl130_54 | spl130_55), inference(avatar_split_clause, [], [f522, f1427, f1423, f1419, f1415, f1411])).
fof(f522, plain, ((e4 = op(e2, e4)) | (e3 = op(e2, e4)) | (e2 = op(e2, e4)) | (e1 = op(e2, e4)) | (e0 = op(e2, e4))), inference(cnf_transformation, [], [f1])).
fof(f1409, plain, (spl130_46 | spl130_47 | spl130_48 | spl130_49 | spl130_50), inference(avatar_split_clause, [], [f523, f1406, f1402, f1398, f1394, f1390])).
fof(f523, plain, ((e4 = op(e3, e0)) | (e3 = op(e3, e0)) | (e2 = op(e3, e0)) | (e1 = op(e3, e0)) | (e0 = op(e3, e0))), inference(cnf_transformation, [], [f1])).
fof(f1388, plain, (spl130_41 | spl130_42 | spl130_43 | spl130_44 | spl130_45), inference(avatar_split_clause, [], [f524, f1385, f1381, f1377, f1373, f1369])).
fof(f524, plain, ((e4 = op(e3, e1)) | (e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (e0 = op(e3, e1))), inference(cnf_transformation, [], [f1])).
fof(f1367, plain, (spl130_36 | spl130_37 | spl130_38 | spl130_39 | spl130_40), inference(avatar_split_clause, [], [f525, f1364, f1360, f1356, f1352, f1348])).
fof(f525, plain, ((e4 = op(e3, e2)) | (e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (e0 = op(e3, e2))), inference(cnf_transformation, [], [f1])).
fof(f1346, plain, (spl130_31 | spl130_32 | spl130_33 | spl130_34 | spl130_35), inference(avatar_split_clause, [], [f526, f1343, f1339, f1335, f1331, f1327])).
fof(f526, plain, ((e4 = op(e3, e3)) | (e3 = op(e3, e3)) | (e2 = op(e3, e3)) | (e1 = op(e3, e3)) | (e0 = op(e3, e3))), inference(cnf_transformation, [], [f1])).
fof(f1325, plain, (spl130_26 | spl130_27 | spl130_28 | spl130_29 | spl130_30), inference(avatar_split_clause, [], [f527, f1322, f1318, f1314, f1310, f1306])).
fof(f527, plain, ((e4 = op(e3, e4)) | (e3 = op(e3, e4)) | (e2 = op(e3, e4)) | (e1 = op(e3, e4)) | (e0 = op(e3, e4))), inference(cnf_transformation, [], [f1])).
fof(f1304, plain, (spl130_21 | spl130_22 | spl130_23 | spl130_24 | spl130_25), inference(avatar_split_clause, [], [f528, f1301, f1297, f1293, f1289, f1285])).
fof(f528, plain, ((e4 = op(e4, e0)) | (e3 = op(e4, e0)) | (e2 = op(e4, e0)) | (e1 = op(e4, e0)) | (e0 = op(e4, e0))), inference(cnf_transformation, [], [f1])).
fof(f1262, plain, (spl130_11 | spl130_12 | spl130_13 | spl130_14 | spl130_15), inference(avatar_split_clause, [], [f530, f1259, f1255, f1251, f1247, f1243])).
fof(f530, plain, ((e4 = op(e4, e2)) | (e3 = op(e4, e2)) | (e2 = op(e4, e2)) | (e1 = op(e4, e2)) | (e0 = op(e4, e2))), inference(cnf_transformation, [], [f1])).
fof(f1241, plain, (spl130_6 | spl130_7 | spl130_8 | spl130_9 | spl130_10), inference(avatar_split_clause, [], [f531, f1238, f1234, f1230, f1226, f1222])).
fof(f531, plain, ((e4 = op(e4, e3)) | (e3 = op(e4, e3)) | (e2 = op(e4, e3)) | (e1 = op(e4, e3)) | (e0 = op(e4, e3))), inference(cnf_transformation, [], [f1])).
fof(f1220, plain, (spl130_1 | spl130_2 | spl130_3 | spl130_4 | spl130_5), inference(avatar_split_clause, [], [f532, f1217, f1213, f1209, f1205, f1201])).
fof(f532, plain, ((e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | (e1 = op(e4, e4)) | (e0 = op(e4, e4))), inference(cnf_transformation, [], [f1])).