fof(f814, plain, $false, inference(subsumption_resolution, [], [f813, f588])).
fof(f588, plain, ~ sP0, inference(subsumption_resolution, [], [f587, f341])).
fof(f341, plain, (unit = op(unit, unit)), inference(definition_unfolding, [], [f112, f128, f128, f128])).
fof(f128, plain, (e0 = unit), inference(cnf_transformation, [], [f3])).
fof(f3, plain, (e0 = unit), file('/home/ubuntu/library/tptp/Problems/ALG/ALG043+1.p', ax3)).
fof(f112, plain, (e0 = op(e0, e0)), inference(cnf_transformation, [], [f2])).
fof(f2, plain, ((e0 = op(e3, e3)) & (e1 = op(e3, e2)) & (e2 = op(e3, e1)) & (e3 = op(e3, e0)) & (e1 = op(e2, e3)) & (e0 = op(e2, e2)) & (e3 = op(e2, e1)) & (e2 = op(e2, e0)) & (e2 = op(e1, e3)) & (e3 = op(e1, e2)) & (e0 = op(e1, e1)) & (e1 = op(e1, e0)) & (e3 = op(e0, e3)) & (e2 = op(e0, e2)) & (e1 = op(e0, e1)) & (e0 = op(e0, e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG043+1.p', ax2)).
fof(f587, plain, (~ (unit = op(unit, unit)) | ~ sP0), inference(subsumption_resolution, [], [f586, f336])).
fof(f336, plain, (op(e1, e1) = unit), inference(definition_unfolding, [], [f117, f128])).
fof(f117, plain, (e0 = op(e1, e1)), inference(cnf_transformation, [], [f2])).
fof(f586, plain, (~ (op(e1, e1) = unit) | ~ (unit = op(unit, unit)) | ~ sP0), inference(subsumption_resolution, [], [f585, f334])).
fof(f334, plain, (op(e2, e2) = unit), inference(definition_unfolding, [], [f122, f128])).
fof(f122, plain, (e0 = op(e2, e2)), inference(cnf_transformation, [], [f2])).
fof(f585, plain, (~ (op(e2, e2) = unit) | ~ (op(e1, e1) = unit) | ~ (unit = op(unit, unit)) | ~ sP0), inference(subsumption_resolution, [], [f456, f332])).
fof(f332, plain, (op(e3, e3) = unit), inference(definition_unfolding, [], [f127, f128])).
fof(f127, plain, (e0 = op(e3, e3)), inference(cnf_transformation, [], [f2])).
fof(f456, plain, (~ (op(e3, e3) = unit) | ~ (op(e2, e2) = unit) | ~ (op(e1, e1) = unit) | ~ (unit = op(unit, unit)) | ~ sP0), inference(definition_unfolding, [], [f321, f128, f128, f128, f128, f128, f128])).
fof(f321, plain, (~ (e0 = op(e3, e3)) | ~ (e0 = op(e2, e2)) | ~ (e0 = op(e1, e1)) | ~ (e0 = op(e0, e0)) | ~ sP0), inference(cnf_transformation, [], [f105])).
fof(f105, plain, (((~ (e3 = op(e3, e3)) | ~ (e3 = op(e2, e2)) | ~ (e3 = op(e1, e1)) | ~ (e3 = op(e0, e0))) & (~ (e2 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e2 = op(e1, e1)) | ~ (e2 = op(e0, e0))) & (~ (e1 = op(e3, e3)) | ~ (e1 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e1 = op(e0, e0))) & (~ (e0 = op(e3, e3)) | ~ (e0 = op(e2, e2)) | ~ (e0 = op(e1, e1)) | ~ (e0 = op(e0, e0)))) | ~ sP0), inference(nnf_transformation, [], [f7])).
fof(f7, plain, (((~ (e3 = op(e3, e3)) | ~ (e3 = op(e2, e2)) | ~ (e3 = op(e1, e1)) | ~ (e3 = op(e0, e0))) & (~ (e2 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e2 = op(e1, e1)) | ~ (e2 = op(e0, e0))) & (~ (e1 = op(e3, e3)) | ~ (e1 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e1 = op(e0, e0))) & (~ (e0 = op(e3, e3)) | ~ (e0 = op(e2, e2)) | ~ (e0 = op(e1, e1)) | ~ (e0 = op(e0, e0)))) | ~ sP0), inference(usedef, [], [e7])).
fof(e7, plain, (sP0 <=> ((~ (e3 = op(e3, e3)) | ~ (e3 = op(e2, e2)) | ~ (e3 = op(e1, e1)) | ~ (e3 = op(e0, e0))) & (~ (e2 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e2 = op(e1, e1)) | ~ (e2 = op(e0, e0))) & (~ (e1 = op(e3, e3)) | ~ (e1 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e1 = op(e0, e0))) & (~ (e0 = op(e3, e3)) | ~ (e0 = op(e2, e2)) | ~ (e0 = op(e1, e1)) | ~ (e0 = op(e0, e0))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f813, plain, sP0, inference(subsumption_resolution, [], [f812, f581])).
fof(f581, plain, ~ sP1, inference(subsumption_resolution, [], [f452, f341])).
fof(f452, plain, (~ (unit = op(unit, unit)) | ~ sP1), inference(definition_unfolding, [], [f317, f128, f128, f128])).
fof(f317, plain, (~ (e0 = op(e0, e0)) | ~ sP1), inference(cnf_transformation, [], [f104])).
fof(f104, plain, ((~ (e3 = op(e0, e0)) & ~ (e2 = op(e0, e0)) & ~ (e1 = op(e0, e0)) & ~ (e0 = op(e0, e0))) | ~ sP1), inference(nnf_transformation, [], [f8])).
fof(f8, plain, ((~ (e3 = op(e0, e0)) & ~ (e2 = op(e0, e0)) & ~ (e1 = op(e0, e0)) & ~ (e0 = op(e0, e0))) | ~ sP1), inference(usedef, [], [e8])).
fof(e8, plain, (sP1 <=> (~ (e3 = op(e0, e0)) & ~ (e2 = op(e0, e0)) & ~ (e1 = op(e0, e0)) & ~ (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f812, plain, (sP1 | sP0), inference(subsumption_resolution, [], [f811, f577])).
fof(f577, plain, ~ sP2, inference(subsumption_resolution, [], [f447, f340])).
fof(f340, plain, (e1 = op(unit, e1)), inference(definition_unfolding, [], [f113, f128])).
fof(f113, plain, (e1 = op(e0, e1)), inference(cnf_transformation, [], [f2])).
fof(f447, plain, (~ (e1 = op(unit, e1)) | ~ sP2), inference(definition_unfolding, [], [f314, f128])).
fof(f314, plain, (~ (e1 = op(e0, e1)) | ~ sP2), inference(cnf_transformation, [], [f103])).
fof(f103, plain, ((~ (e3 = op(e0, e1)) & ~ (e2 = op(e0, e1)) & ~ (e1 = op(e0, e1)) & ~ (e0 = op(e0, e1))) | ~ sP2), inference(nnf_transformation, [], [f9])).
fof(f9, plain, ((~ (e3 = op(e0, e1)) & ~ (e2 = op(e0, e1)) & ~ (e1 = op(e0, e1)) & ~ (e0 = op(e0, e1))) | ~ sP2), inference(usedef, [], [e9])).
fof(e9, plain, (sP2 <=> (~ (e3 = op(e0, e1)) & ~ (e2 = op(e0, e1)) & ~ (e1 = op(e0, e1)) & ~ (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f811, plain, (sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f810, f574])).
fof(f574, plain, ~ sP3, inference(subsumption_resolution, [], [f442, f339])).
fof(f339, plain, (e2 = op(unit, e2)), inference(definition_unfolding, [], [f114, f128])).
fof(f114, plain, (e2 = op(e0, e2)), inference(cnf_transformation, [], [f2])).
fof(f442, plain, (~ (e2 = op(unit, e2)) | ~ sP3), inference(definition_unfolding, [], [f311, f128])).
fof(f311, plain, (~ (e2 = op(e0, e2)) | ~ sP3), inference(cnf_transformation, [], [f102])).
fof(f102, plain, ((~ (e3 = op(e0, e2)) & ~ (e2 = op(e0, e2)) & ~ (e1 = op(e0, e2)) & ~ (e0 = op(e0, e2))) | ~ sP3), inference(nnf_transformation, [], [f10])).
fof(f10, plain, ((~ (e3 = op(e0, e2)) & ~ (e2 = op(e0, e2)) & ~ (e1 = op(e0, e2)) & ~ (e0 = op(e0, e2))) | ~ sP3), inference(usedef, [], [e10])).
fof(e10, plain, (sP3 <=> (~ (e3 = op(e0, e2)) & ~ (e2 = op(e0, e2)) & ~ (e1 = op(e0, e2)) & ~ (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f810, plain, (sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f809, f572])).
fof(f572, plain, ~ sP4, inference(subsumption_resolution, [], [f437, f338])).
fof(f338, plain, (e3 = op(unit, e3)), inference(definition_unfolding, [], [f115, f128])).
fof(f115, plain, (e3 = op(e0, e3)), inference(cnf_transformation, [], [f2])).
fof(f437, plain, (~ (e3 = op(unit, e3)) | ~ sP4), inference(definition_unfolding, [], [f308, f128])).
fof(f308, plain, (~ (e3 = op(e0, e3)) | ~ sP4), inference(cnf_transformation, [], [f101])).
fof(f101, plain, ((~ (e3 = op(e0, e3)) & ~ (e2 = op(e0, e3)) & ~ (e1 = op(e0, e3)) & ~ (e0 = op(e0, e3))) | ~ sP4), inference(nnf_transformation, [], [f11])).
fof(f11, plain, ((~ (e3 = op(e0, e3)) & ~ (e2 = op(e0, e3)) & ~ (e1 = op(e0, e3)) & ~ (e0 = op(e0, e3))) | ~ sP4), inference(usedef, [], [e11])).
fof(e11, plain, (sP4 <=> (~ (e3 = op(e0, e3)) & ~ (e2 = op(e0, e3)) & ~ (e1 = op(e0, e3)) & ~ (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f809, plain, (sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f808, f571])).
fof(f571, plain, ~ sP5, inference(subsumption_resolution, [], [f435, f337])).
fof(f337, plain, (e1 = op(e1, unit)), inference(definition_unfolding, [], [f116, f128])).
fof(f116, plain, (e1 = op(e1, e0)), inference(cnf_transformation, [], [f2])).
fof(f435, plain, (~ (e1 = op(e1, unit)) | ~ sP5), inference(definition_unfolding, [], [f302, f128])).
fof(f302, plain, (~ (e1 = op(e1, e0)) | ~ sP5), inference(cnf_transformation, [], [f100])).
fof(f100, plain, ((~ (e3 = op(e1, e0)) & ~ (e2 = op(e1, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e1, e0))) | ~ sP5), inference(nnf_transformation, [], [f12])).
fof(f12, plain, ((~ (e3 = op(e1, e0)) & ~ (e2 = op(e1, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e1, e0))) | ~ sP5), inference(usedef, [], [e12])).
fof(e12, plain, (sP5 <=> (~ (e3 = op(e1, e0)) & ~ (e2 = op(e1, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f808, plain, (sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f807, f568])).
fof(f568, plain, ~ sP6, inference(subsumption_resolution, [], [f432, f336])).
fof(f432, plain, (~ (op(e1, e1) = unit) | ~ sP6), inference(definition_unfolding, [], [f297, f128])).
fof(f297, plain, (~ (e0 = op(e1, e1)) | ~ sP6), inference(cnf_transformation, [], [f99])).
fof(f99, plain, ((~ (e3 = op(e1, e1)) & ~ (e2 = op(e1, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e1, e1))) | ~ sP6), inference(nnf_transformation, [], [f13])).
fof(f13, plain, ((~ (e3 = op(e1, e1)) & ~ (e2 = op(e1, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e1, e1))) | ~ sP6), inference(usedef, [], [e13])).
fof(e13, plain, (sP6 <=> (~ (e3 = op(e1, e1)) & ~ (e2 = op(e1, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f807, plain, (sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f806, f564])).
fof(f564, plain, ~ sP7, inference(subsumption_resolution, [], [f296, f118])).
fof(f118, plain, (e3 = op(e1, e2)), inference(cnf_transformation, [], [f2])).
fof(f296, plain, (~ (e3 = op(e1, e2)) | ~ sP7), inference(cnf_transformation, [], [f98])).
fof(f98, plain, ((~ (e3 = op(e1, e2)) & ~ (e2 = op(e1, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e1, e2))) | ~ sP7), inference(nnf_transformation, [], [f14])).
fof(f14, plain, ((~ (e3 = op(e1, e2)) & ~ (e2 = op(e1, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e1, e2))) | ~ sP7), inference(usedef, [], [e14])).
fof(e14, plain, (sP7 <=> (~ (e3 = op(e1, e2)) & ~ (e2 = op(e1, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f806, plain, (sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f805, f563])).
fof(f563, plain, ~ sP8, inference(subsumption_resolution, [], [f291, f119])).
fof(f119, plain, (e2 = op(e1, e3)), inference(cnf_transformation, [], [f2])).
fof(f291, plain, (~ (e2 = op(e1, e3)) | ~ sP8), inference(cnf_transformation, [], [f97])).
fof(f97, plain, ((~ (e3 = op(e1, e3)) & ~ (e2 = op(e1, e3)) & ~ (e1 = op(e1, e3)) & ~ (e0 = op(e1, e3))) | ~ sP8), inference(nnf_transformation, [], [f15])).
fof(f15, plain, ((~ (e3 = op(e1, e3)) & ~ (e2 = op(e1, e3)) & ~ (e1 = op(e1, e3)) & ~ (e0 = op(e1, e3))) | ~ sP8), inference(usedef, [], [e15])).
fof(e15, plain, (sP8 <=> (~ (e3 = op(e1, e3)) & ~ (e2 = op(e1, e3)) & ~ (e1 = op(e1, e3)) & ~ (e0 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f805, plain, (sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f804, f561])).
fof(f561, plain, ~ sP9, inference(subsumption_resolution, [], [f427, f335])).
fof(f335, plain, (e2 = op(e2, unit)), inference(definition_unfolding, [], [f120, f128])).
fof(f120, plain, (e2 = op(e2, e0)), inference(cnf_transformation, [], [f2])).
fof(f427, plain, (~ (e2 = op(e2, unit)) | ~ sP9), inference(definition_unfolding, [], [f287, f128])).
fof(f287, plain, (~ (e2 = op(e2, e0)) | ~ sP9), inference(cnf_transformation, [], [f96])).
fof(f96, plain, ((~ (e3 = op(e2, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e2, e0)) & ~ (e0 = op(e2, e0))) | ~ sP9), inference(nnf_transformation, [], [f16])).
fof(f16, plain, ((~ (e3 = op(e2, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e2, e0)) & ~ (e0 = op(e2, e0))) | ~ sP9), inference(usedef, [], [e16])).
fof(e16, plain, (sP9 <=> (~ (e3 = op(e2, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e2, e0)) & ~ (e0 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f804, plain, (sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f803, f559])).
fof(f559, plain, ~ sP10, inference(subsumption_resolution, [], [f284, f121])).
fof(f121, plain, (e3 = op(e2, e1)), inference(cnf_transformation, [], [f2])).
fof(f284, plain, (~ (e3 = op(e2, e1)) | ~ sP10), inference(cnf_transformation, [], [f95])).
fof(f95, plain, ((~ (e3 = op(e2, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e2, e1)) & ~ (e0 = op(e2, e1))) | ~ sP10), inference(nnf_transformation, [], [f17])).
fof(f17, plain, ((~ (e3 = op(e2, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e2, e1)) & ~ (e0 = op(e2, e1))) | ~ sP10), inference(usedef, [], [e17])).
fof(e17, plain, (sP10 <=> (~ (e3 = op(e2, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e2, e1)) & ~ (e0 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f803, plain, (sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f802, f558])).
fof(f558, plain, ~ sP11, inference(subsumption_resolution, [], [f424, f334])).
fof(f424, plain, (~ (op(e2, e2) = unit) | ~ sP11), inference(definition_unfolding, [], [f277, f128])).
fof(f277, plain, (~ (e0 = op(e2, e2)) | ~ sP11), inference(cnf_transformation, [], [f94])).
fof(f94, plain, ((~ (e3 = op(e2, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e2, e2)) & ~ (e0 = op(e2, e2))) | ~ sP11), inference(nnf_transformation, [], [f18])).
fof(f18, plain, ((~ (e3 = op(e2, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e2, e2)) & ~ (e0 = op(e2, e2))) | ~ sP11), inference(usedef, [], [e18])).
fof(e18, plain, (sP11 <=> (~ (e3 = op(e2, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e2, e2)) & ~ (e0 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f802, plain, (sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f801, f554])).
fof(f554, plain, ~ sP12, inference(subsumption_resolution, [], [f274, f123])).
fof(f123, plain, (e1 = op(e2, e3)), inference(cnf_transformation, [], [f2])).
fof(f274, plain, (~ (e1 = op(e2, e3)) | ~ sP12), inference(cnf_transformation, [], [f93])).
fof(f93, plain, ((~ (e3 = op(e2, e3)) & ~ (e2 = op(e2, e3)) & ~ (e1 = op(e2, e3)) & ~ (e0 = op(e2, e3))) | ~ sP12), inference(nnf_transformation, [], [f19])).
fof(f19, plain, ((~ (e3 = op(e2, e3)) & ~ (e2 = op(e2, e3)) & ~ (e1 = op(e2, e3)) & ~ (e0 = op(e2, e3))) | ~ sP12), inference(usedef, [], [e19])).
fof(e19, plain, (sP12 <=> (~ (e3 = op(e2, e3)) & ~ (e2 = op(e2, e3)) & ~ (e1 = op(e2, e3)) & ~ (e0 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f801, plain, (sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f800, f551])).
fof(f551, plain, ~ sP13, inference(subsumption_resolution, [], [f419, f333])).
fof(f333, plain, (e3 = op(e3, unit)), inference(definition_unfolding, [], [f124, f128])).
fof(f124, plain, (e3 = op(e3, e0)), inference(cnf_transformation, [], [f2])).
fof(f419, plain, (~ (e3 = op(e3, unit)) | ~ sP13), inference(definition_unfolding, [], [f272, f128])).
fof(f272, plain, (~ (e3 = op(e3, e0)) | ~ sP13), inference(cnf_transformation, [], [f92])).
fof(f92, plain, ((~ (e3 = op(e3, e0)) & ~ (e2 = op(e3, e0)) & ~ (e1 = op(e3, e0)) & ~ (e0 = op(e3, e0))) | ~ sP13), inference(nnf_transformation, [], [f20])).
fof(f20, plain, ((~ (e3 = op(e3, e0)) & ~ (e2 = op(e3, e0)) & ~ (e1 = op(e3, e0)) & ~ (e0 = op(e3, e0))) | ~ sP13), inference(usedef, [], [e20])).
fof(e20, plain, (sP13 <=> (~ (e3 = op(e3, e0)) & ~ (e2 = op(e3, e0)) & ~ (e1 = op(e3, e0)) & ~ (e0 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f800, plain, (sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f799, f550])).
fof(f550, plain, ~ sP14, inference(subsumption_resolution, [], [f267, f125])).
fof(f125, plain, (e2 = op(e3, e1)), inference(cnf_transformation, [], [f2])).
fof(f267, plain, (~ (e2 = op(e3, e1)) | ~ sP14), inference(cnf_transformation, [], [f91])).
fof(f91, plain, ((~ (e3 = op(e3, e1)) & ~ (e2 = op(e3, e1)) & ~ (e1 = op(e3, e1)) & ~ (e0 = op(e3, e1))) | ~ sP14), inference(nnf_transformation, [], [f21])).
fof(f21, plain, ((~ (e3 = op(e3, e1)) & ~ (e2 = op(e3, e1)) & ~ (e1 = op(e3, e1)) & ~ (e0 = op(e3, e1))) | ~ sP14), inference(usedef, [], [e21])).
fof(e21, plain, (sP14 <=> (~ (e3 = op(e3, e1)) & ~ (e2 = op(e3, e1)) & ~ (e1 = op(e3, e1)) & ~ (e0 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f799, plain, (sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f798, f548])).
fof(f548, plain, ~ sP15, inference(subsumption_resolution, [], [f262, f126])).
fof(f126, plain, (e1 = op(e3, e2)), inference(cnf_transformation, [], [f2])).
fof(f262, plain, (~ (e1 = op(e3, e2)) | ~ sP15), inference(cnf_transformation, [], [f90])).
fof(f90, plain, ((~ (e3 = op(e3, e2)) & ~ (e2 = op(e3, e2)) & ~ (e1 = op(e3, e2)) & ~ (e0 = op(e3, e2))) | ~ sP15), inference(nnf_transformation, [], [f22])).
fof(f22, plain, ((~ (e3 = op(e3, e2)) & ~ (e2 = op(e3, e2)) & ~ (e1 = op(e3, e2)) & ~ (e0 = op(e3, e2))) | ~ sP15), inference(usedef, [], [e22])).
fof(e22, plain, (sP15 <=> (~ (e3 = op(e3, e2)) & ~ (e2 = op(e3, e2)) & ~ (e1 = op(e3, e2)) & ~ (e0 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f798, plain, (sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f797, f545])).
fof(f545, plain, ~ sP16, inference(subsumption_resolution, [], [f416, f332])).
fof(f416, plain, (~ (op(e3, e3) = unit) | ~ sP16), inference(definition_unfolding, [], [f257, f128])).
fof(f257, plain, (~ (e0 = op(e3, e3)) | ~ sP16), inference(cnf_transformation, [], [f89])).
fof(f89, plain, ((~ (e3 = op(e3, e3)) & ~ (e2 = op(e3, e3)) & ~ (e1 = op(e3, e3)) & ~ (e0 = op(e3, e3))) | ~ sP16), inference(nnf_transformation, [], [f23])).
fof(f23, plain, ((~ (e3 = op(e3, e3)) & ~ (e2 = op(e3, e3)) & ~ (e1 = op(e3, e3)) & ~ (e0 = op(e3, e3))) | ~ sP16), inference(usedef, [], [e23])).
fof(e23, plain, (sP16 <=> (~ (e3 = op(e3, e3)) & ~ (e2 = op(e3, e3)) & ~ (e1 = op(e3, e3)) & ~ (e0 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f797, plain, (sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f796, f341])).
fof(f796, plain, (~ (unit = op(unit, unit)) | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f795, f340])).
fof(f795, plain, (~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f794, f337])).
fof(f794, plain, (~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f793, f339])).
fof(f793, plain, (~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f792, f335])).
fof(f792, plain, (~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f791, f333])).
fof(f791, plain, (~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f790, f465])).
fof(f465, plain, ~ sP17, inference(trivial_inequality_removal, [], [f415])).
fof(f415, plain, (~ (unit = unit) | ~ sP17), inference(definition_unfolding, [], [f253, f128])).
fof(f253, plain, (~ (e0 = unit) | ~ sP17), inference(cnf_transformation, [], [f88])).
fof(f88, plain, ((~ (e3 = unit) & ~ (e2 = unit) & ~ (e1 = unit) & ~ (e0 = unit)) | ~ sP17), inference(nnf_transformation, [], [f24])).
fof(f24, plain, ((~ (e3 = unit) & ~ (e2 = unit) & ~ (e1 = unit) & ~ (e0 = unit)) | ~ sP17), inference(usedef, [], [e24])).
fof(e24, plain, (sP17 <=> (~ (e3 = unit) & ~ (e2 = unit) & ~ (e1 = unit) & ~ (e0 = unit))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f790, plain, (sP17 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f789, f541])).
fof(f541, plain, ~ sP18, inference(subsumption_resolution, [], [f414, f341])).
fof(f414, plain, (~ (unit = op(unit, unit)) | ~ sP18), inference(definition_unfolding, [], [f249, f128, f128, f128])).
fof(f249, plain, (~ (e0 = op(e0, e0)) | ~ sP18), inference(cnf_transformation, [], [f87])).
fof(f87, plain, ((~ (e0 = op(e0, e3)) & ~ (e0 = op(e0, e2)) & ~ (e0 = op(e0, e1)) & ~ (e0 = op(e0, e0))) | ~ sP18), inference(nnf_transformation, [], [f25])).
fof(f25, plain, ((~ (e0 = op(e0, e3)) & ~ (e0 = op(e0, e2)) & ~ (e0 = op(e0, e1)) & ~ (e0 = op(e0, e0))) | ~ sP18), inference(usedef, [], [e25])).
fof(e25, plain, (sP18 <=> (~ (e0 = op(e0, e3)) & ~ (e0 = op(e0, e2)) & ~ (e0 = op(e0, e1)) & ~ (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f789, plain, (sP18 | sP17 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f788, f537])).
fof(f537, plain, ~ sP19, inference(subsumption_resolution, [], [f410, f341])).
fof(f410, plain, (~ (unit = op(unit, unit)) | ~ sP19), inference(definition_unfolding, [], [f245, f128, f128, f128])).
fof(f245, plain, (~ (e0 = op(e0, e0)) | ~ sP19), inference(cnf_transformation, [], [f86])).
fof(f86, plain, ((~ (e0 = op(e3, e0)) & ~ (e0 = op(e2, e0)) & ~ (e0 = op(e1, e0)) & ~ (e0 = op(e0, e0))) | ~ sP19), inference(nnf_transformation, [], [f26])).
fof(f26, plain, ((~ (e0 = op(e3, e0)) & ~ (e0 = op(e2, e0)) & ~ (e0 = op(e1, e0)) & ~ (e0 = op(e0, e0))) | ~ sP19), inference(usedef, [], [e26])).
fof(e26, plain, (sP19 <=> (~ (e0 = op(e3, e0)) & ~ (e0 = op(e2, e0)) & ~ (e0 = op(e1, e0)) & ~ (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP19])])).
fof(f788, plain, (sP19 | sP18 | sP17 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f787, f533])).
fof(f533, plain, ~ sP20, inference(subsumption_resolution, [], [f405, f340])).
fof(f405, plain, (~ (e1 = op(unit, e1)) | ~ sP20), inference(definition_unfolding, [], [f242, f128])).
fof(f242, plain, (~ (e1 = op(e0, e1)) | ~ sP20), inference(cnf_transformation, [], [f85])).
fof(f85, plain, ((~ (e1 = op(e0, e3)) & ~ (e1 = op(e0, e2)) & ~ (e1 = op(e0, e1)) & ~ (e1 = op(e0, e0))) | ~ sP20), inference(nnf_transformation, [], [f27])).
fof(f27, plain, ((~ (e1 = op(e0, e3)) & ~ (e1 = op(e0, e2)) & ~ (e1 = op(e0, e1)) & ~ (e1 = op(e0, e0))) | ~ sP20), inference(usedef, [], [e27])).
fof(e27, plain, (sP20 <=> (~ (e1 = op(e0, e3)) & ~ (e1 = op(e0, e2)) & ~ (e1 = op(e0, e1)) & ~ (e1 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP20])])).
fof(f787, plain, (sP20 | sP19 | sP18 | sP17 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f786, f530])).
fof(f530, plain, ~ sP21, inference(subsumption_resolution, [], [f401, f337])).
fof(f401, plain, (~ (e1 = op(e1, unit)) | ~ sP21), inference(definition_unfolding, [], [f238, f128])).
fof(f238, plain, (~ (e1 = op(e1, e0)) | ~ sP21), inference(cnf_transformation, [], [f84])).
fof(f84, plain, ((~ (e1 = op(e3, e0)) & ~ (e1 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e1 = op(e0, e0))) | ~ sP21), inference(nnf_transformation, [], [f28])).
fof(f28, plain, ((~ (e1 = op(e3, e0)) & ~ (e1 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e1 = op(e0, e0))) | ~ sP21), inference(usedef, [], [e28])).
fof(e28, plain, (sP21 <=> (~ (e1 = op(e3, e0)) & ~ (e1 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e1 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP21])])).
fof(f786, plain, (sP21 | sP20 | sP19 | sP18 | sP17 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f785, f527])).
fof(f527, plain, ~ sP22, inference(subsumption_resolution, [], [f396, f339])).
fof(f396, plain, (~ (e2 = op(unit, e2)) | ~ sP22), inference(definition_unfolding, [], [f235, f128])).
fof(f235, plain, (~ (e2 = op(e0, e2)) | ~ sP22), inference(cnf_transformation, [], [f83])).
fof(f83, plain, ((~ (e2 = op(e0, e3)) & ~ (e2 = op(e0, e2)) & ~ (e2 = op(e0, e1)) & ~ (e2 = op(e0, e0))) | ~ sP22), inference(nnf_transformation, [], [f29])).
fof(f29, plain, ((~ (e2 = op(e0, e3)) & ~ (e2 = op(e0, e2)) & ~ (e2 = op(e0, e1)) & ~ (e2 = op(e0, e0))) | ~ sP22), inference(usedef, [], [e29])).
fof(e29, plain, (sP22 <=> (~ (e2 = op(e0, e3)) & ~ (e2 = op(e0, e2)) & ~ (e2 = op(e0, e1)) & ~ (e2 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP22])])).
fof(f785, plain, (sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f784, f525])).
fof(f525, plain, ~ sP23, inference(subsumption_resolution, [], [f392, f335])).
fof(f392, plain, (~ (e2 = op(e2, unit)) | ~ sP23), inference(definition_unfolding, [], [f231, f128])).
fof(f231, plain, (~ (e2 = op(e2, e0)) | ~ sP23), inference(cnf_transformation, [], [f82])).
fof(f82, plain, ((~ (e2 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e2 = op(e1, e0)) & ~ (e2 = op(e0, e0))) | ~ sP23), inference(nnf_transformation, [], [f30])).
fof(f30, plain, ((~ (e2 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e2 = op(e1, e0)) & ~ (e2 = op(e0, e0))) | ~ sP23), inference(usedef, [], [e30])).
fof(e30, plain, (sP23 <=> (~ (e2 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e2 = op(e1, e0)) & ~ (e2 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP23])])).
fof(f784, plain, (sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f783, f523])).
fof(f523, plain, ~ sP24, inference(subsumption_resolution, [], [f387, f338])).
fof(f387, plain, (~ (e3 = op(unit, e3)) | ~ sP24), inference(definition_unfolding, [], [f228, f128])).
fof(f228, plain, (~ (e3 = op(e0, e3)) | ~ sP24), inference(cnf_transformation, [], [f81])).
fof(f81, plain, ((~ (e3 = op(e0, e3)) & ~ (e3 = op(e0, e2)) & ~ (e3 = op(e0, e1)) & ~ (e3 = op(e0, e0))) | ~ sP24), inference(nnf_transformation, [], [f31])).
fof(f31, plain, ((~ (e3 = op(e0, e3)) & ~ (e3 = op(e0, e2)) & ~ (e3 = op(e0, e1)) & ~ (e3 = op(e0, e0))) | ~ sP24), inference(usedef, [], [e31])).
fof(e31, plain, (sP24 <=> (~ (e3 = op(e0, e3)) & ~ (e3 = op(e0, e2)) & ~ (e3 = op(e0, e1)) & ~ (e3 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP24])])).
fof(f783, plain, (sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f782, f522])).
fof(f522, plain, ~ sP25, inference(subsumption_resolution, [], [f383, f333])).
fof(f383, plain, (~ (e3 = op(e3, unit)) | ~ sP25), inference(definition_unfolding, [], [f224, f128])).
fof(f224, plain, (~ (e3 = op(e3, e0)) | ~ sP25), inference(cnf_transformation, [], [f80])).
fof(f80, plain, ((~ (e3 = op(e3, e0)) & ~ (e3 = op(e2, e0)) & ~ (e3 = op(e1, e0)) & ~ (e3 = op(e0, e0))) | ~ sP25), inference(nnf_transformation, [], [f32])).
fof(f32, plain, ((~ (e3 = op(e3, e0)) & ~ (e3 = op(e2, e0)) & ~ (e3 = op(e1, e0)) & ~ (e3 = op(e0, e0))) | ~ sP25), inference(usedef, [], [e32])).
fof(e32, plain, (sP25 <=> (~ (e3 = op(e3, e0)) & ~ (e3 = op(e2, e0)) & ~ (e3 = op(e1, e0)) & ~ (e3 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP25])])).
fof(f782, plain, (sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f781, f521])).
fof(f521, plain, ~ sP26, inference(subsumption_resolution, [], [f381, f336])).
fof(f381, plain, (~ (op(e1, e1) = unit) | ~ sP26), inference(definition_unfolding, [], [f218, f128])).
fof(f218, plain, (~ (e0 = op(e1, e1)) | ~ sP26), inference(cnf_transformation, [], [f79])).
fof(f79, plain, ((~ (e0 = op(e1, e3)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e1, e0))) | ~ sP26), inference(nnf_transformation, [], [f33])).
fof(f33, plain, ((~ (e0 = op(e1, e3)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e1, e0))) | ~ sP26), inference(usedef, [], [e33])).
fof(e33, plain, (sP26 <=> (~ (e0 = op(e1, e3)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP26])])).
fof(f781, plain, (sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f780, f518])).
fof(f518, plain, ~ sP27, inference(subsumption_resolution, [], [f377, f336])).
fof(f377, plain, (~ (op(e1, e1) = unit) | ~ sP27), inference(definition_unfolding, [], [f214, f128])).
fof(f214, plain, (~ (e0 = op(e1, e1)) | ~ sP27), inference(cnf_transformation, [], [f78])).
fof(f78, plain, ((~ (e0 = op(e3, e1)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e0, e1))) | ~ sP27), inference(nnf_transformation, [], [f34])).
fof(f34, plain, ((~ (e0 = op(e3, e1)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e0, e1))) | ~ sP27), inference(usedef, [], [e34])).
fof(e34, plain, (sP27 <=> (~ (e0 = op(e3, e1)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP27])])).
fof(f780, plain, (sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f779, f515])).
fof(f515, plain, ~ sP28, inference(subsumption_resolution, [], [f374, f337])).
fof(f374, plain, (~ (e1 = op(e1, unit)) | ~ sP28), inference(definition_unfolding, [], [f209, f128])).
fof(f209, plain, (~ (e1 = op(e1, e0)) | ~ sP28), inference(cnf_transformation, [], [f77])).
fof(f77, plain, ((~ (e1 = op(e1, e3)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e1, e0))) | ~ sP28), inference(nnf_transformation, [], [f35])).
fof(f35, plain, ((~ (e1 = op(e1, e3)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e1, e0))) | ~ sP28), inference(usedef, [], [e35])).
fof(e35, plain, (sP28 <=> (~ (e1 = op(e1, e3)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP28])])).
fof(f779, plain, (sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f778, f511])).
fof(f511, plain, ~ sP29, inference(subsumption_resolution, [], [f373, f340])).
fof(f373, plain, (~ (e1 = op(unit, e1)) | ~ sP29), inference(definition_unfolding, [], [f205, f128])).
fof(f205, plain, (~ (e1 = op(e0, e1)) | ~ sP29), inference(cnf_transformation, [], [f76])).
fof(f76, plain, ((~ (e1 = op(e3, e1)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e0, e1))) | ~ sP29), inference(nnf_transformation, [], [f36])).
fof(f36, plain, ((~ (e1 = op(e3, e1)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e0, e1))) | ~ sP29), inference(usedef, [], [e36])).
fof(e36, plain, (sP29 <=> (~ (e1 = op(e3, e1)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP29])])).
fof(f778, plain, (sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f777, f507])).
fof(f507, plain, ~ sP30, inference(subsumption_resolution, [], [f204, f119])).
fof(f204, plain, (~ (e2 = op(e1, e3)) | ~ sP30), inference(cnf_transformation, [], [f75])).
fof(f75, plain, ((~ (e2 = op(e1, e3)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e1, e0))) | ~ sP30), inference(nnf_transformation, [], [f37])).
fof(f37, plain, ((~ (e2 = op(e1, e3)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e1, e0))) | ~ sP30), inference(usedef, [], [e37])).
fof(e37, plain, (sP30 <=> (~ (e2 = op(e1, e3)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP30])])).
fof(f777, plain, (sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f776, f506])).
fof(f506, plain, ~ sP31, inference(subsumption_resolution, [], [f200, f125])).
fof(f200, plain, (~ (e2 = op(e3, e1)) | ~ sP31), inference(cnf_transformation, [], [f74])).
fof(f74, plain, ((~ (e2 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e0, e1))) | ~ sP31), inference(nnf_transformation, [], [f38])).
fof(f38, plain, ((~ (e2 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e0, e1))) | ~ sP31), inference(usedef, [], [e38])).
fof(e38, plain, (sP31 <=> (~ (e2 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP31])])).
fof(f776, plain, (sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f775, f505])).
fof(f505, plain, ~ sP32, inference(subsumption_resolution, [], [f195, f118])).
fof(f195, plain, (~ (e3 = op(e1, e2)) | ~ sP32), inference(cnf_transformation, [], [f73])).
fof(f73, plain, ((~ (e3 = op(e1, e3)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e1, e0))) | ~ sP32), inference(nnf_transformation, [], [f39])).
fof(f39, plain, ((~ (e3 = op(e1, e3)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e1, e0))) | ~ sP32), inference(usedef, [], [e39])).
fof(e39, plain, (sP32 <=> (~ (e3 = op(e1, e3)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP32])])).
fof(f775, plain, (sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f774, f503])).
fof(f503, plain, ~ sP33, inference(subsumption_resolution, [], [f191, f121])).
fof(f191, plain, (~ (e3 = op(e2, e1)) | ~ sP33), inference(cnf_transformation, [], [f72])).
fof(f72, plain, ((~ (e3 = op(e3, e1)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e0, e1))) | ~ sP33), inference(nnf_transformation, [], [f40])).
fof(f40, plain, ((~ (e3 = op(e3, e1)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e0, e1))) | ~ sP33), inference(usedef, [], [e40])).
fof(e40, plain, (sP33 <=> (~ (e3 = op(e3, e1)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP33])])).
fof(f774, plain, (sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f773, f501])).
fof(f501, plain, ~ sP34, inference(subsumption_resolution, [], [f366, f334])).
fof(f366, plain, (~ (op(e2, e2) = unit) | ~ sP34), inference(definition_unfolding, [], [f187, f128])).
fof(f187, plain, (~ (e0 = op(e2, e2)) | ~ sP34), inference(cnf_transformation, [], [f71])).
fof(f71, plain, ((~ (e0 = op(e2, e3)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e2, e0))) | ~ sP34), inference(nnf_transformation, [], [f41])).
fof(f41, plain, ((~ (e0 = op(e2, e3)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e2, e0))) | ~ sP34), inference(usedef, [], [e41])).
fof(e41, plain, (sP34 <=> (~ (e0 = op(e2, e3)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP34])])).
fof(f773, plain, (sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f772, f499])).
fof(f499, plain, ~ sP35, inference(subsumption_resolution, [], [f362, f334])).
fof(f362, plain, (~ (op(e2, e2) = unit) | ~ sP35), inference(definition_unfolding, [], [f183, f128])).
fof(f183, plain, (~ (e0 = op(e2, e2)) | ~ sP35), inference(cnf_transformation, [], [f70])).
fof(f70, plain, ((~ (e0 = op(e3, e2)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e0, e2))) | ~ sP35), inference(nnf_transformation, [], [f42])).
fof(f42, plain, ((~ (e0 = op(e3, e2)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e0, e2))) | ~ sP35), inference(usedef, [], [e42])).
fof(e42, plain, (sP35 <=> (~ (e0 = op(e3, e2)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP35])])).
fof(f772, plain, (sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f771, f497])).
fof(f497, plain, ~ sP36, inference(subsumption_resolution, [], [f180, f123])).
fof(f180, plain, (~ (e1 = op(e2, e3)) | ~ sP36), inference(cnf_transformation, [], [f69])).
fof(f69, plain, ((~ (e1 = op(e2, e3)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e2, e0))) | ~ sP36), inference(nnf_transformation, [], [f43])).
fof(f43, plain, ((~ (e1 = op(e2, e3)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e2, e0))) | ~ sP36), inference(usedef, [], [e43])).
fof(e43, plain, (sP36 <=> (~ (e1 = op(e2, e3)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP36])])).
fof(f771, plain, (sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f770, f496])).
fof(f496, plain, ~ sP37, inference(subsumption_resolution, [], [f176, f126])).
fof(f176, plain, (~ (e1 = op(e3, e2)) | ~ sP37), inference(cnf_transformation, [], [f68])).
fof(f68, plain, ((~ (e1 = op(e3, e2)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e0, e2))) | ~ sP37), inference(nnf_transformation, [], [f44])).
fof(f44, plain, ((~ (e1 = op(e3, e2)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e0, e2))) | ~ sP37), inference(usedef, [], [e44])).
fof(e44, plain, (sP37 <=> (~ (e1 = op(e3, e2)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP37])])).
fof(f770, plain, (sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f769, f495])).
fof(f495, plain, ~ sP38, inference(subsumption_resolution, [], [f358, f335])).
fof(f358, plain, (~ (e2 = op(e2, unit)) | ~ sP38), inference(definition_unfolding, [], [f169, f128])).
fof(f169, plain, (~ (e2 = op(e2, e0)) | ~ sP38), inference(cnf_transformation, [], [f67])).
fof(f67, plain, ((~ (e2 = op(e2, e3)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e2, e0))) | ~ sP38), inference(nnf_transformation, [], [f45])).
fof(f45, plain, ((~ (e2 = op(e2, e3)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e2, e0))) | ~ sP38), inference(usedef, [], [e45])).
fof(e45, plain, (sP38 <=> (~ (e2 = op(e2, e3)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP38])])).
fof(f769, plain, (sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f768, f491])).
fof(f491, plain, ~ sP39, inference(subsumption_resolution, [], [f357, f339])).
fof(f357, plain, (~ (e2 = op(unit, e2)) | ~ sP39), inference(definition_unfolding, [], [f165, f128])).
fof(f165, plain, (~ (e2 = op(e0, e2)) | ~ sP39), inference(cnf_transformation, [], [f66])).
fof(f66, plain, ((~ (e2 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e0, e2))) | ~ sP39), inference(nnf_transformation, [], [f46])).
fof(f46, plain, ((~ (e2 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e0, e2))) | ~ sP39), inference(usedef, [], [e46])).
fof(e46, plain, (sP39 <=> (~ (e2 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP39])])).
fof(f768, plain, (sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f767, f487])).
fof(f487, plain, ~ sP40, inference(subsumption_resolution, [], [f162, f121])).
fof(f162, plain, (~ (e3 = op(e2, e1)) | ~ sP40), inference(cnf_transformation, [], [f65])).
fof(f65, plain, ((~ (e3 = op(e2, e3)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e2, e0))) | ~ sP40), inference(nnf_transformation, [], [f47])).
fof(f47, plain, ((~ (e3 = op(e2, e3)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e2, e0))) | ~ sP40), inference(usedef, [], [e47])).
fof(e47, plain, (sP40 <=> (~ (e3 = op(e2, e3)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP40])])).
fof(f767, plain, (sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f766, f484])).
fof(f484, plain, ~ sP41, inference(subsumption_resolution, [], [f158, f118])).
fof(f158, plain, (~ (e3 = op(e1, e2)) | ~ sP41), inference(cnf_transformation, [], [f64])).
fof(f64, plain, ((~ (e3 = op(e3, e2)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e0, e2))) | ~ sP41), inference(nnf_transformation, [], [f48])).
fof(f48, plain, ((~ (e3 = op(e3, e2)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e0, e2))) | ~ sP41), inference(usedef, [], [e48])).
fof(e48, plain, (sP41 <=> (~ (e3 = op(e3, e2)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP41])])).
fof(f766, plain, (sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f765, f481])).
fof(f481, plain, ~ sP42, inference(subsumption_resolution, [], [f351, f332])).
fof(f351, plain, (~ (op(e3, e3) = unit) | ~ sP42), inference(definition_unfolding, [], [f156, f128])).
fof(f156, plain, (~ (e0 = op(e3, e3)) | ~ sP42), inference(cnf_transformation, [], [f63])).
fof(f63, plain, ((~ (e0 = op(e3, e3)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e3, e0))) | ~ sP42), inference(nnf_transformation, [], [f49])).
fof(f49, plain, ((~ (e0 = op(e3, e3)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e3, e0))) | ~ sP42), inference(usedef, [], [e49])).
fof(e49, plain, (sP42 <=> (~ (e0 = op(e3, e3)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP42])])).
fof(f765, plain, (sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f764, f480])).
fof(f480, plain, ~ sP43, inference(subsumption_resolution, [], [f347, f332])).
fof(f347, plain, (~ (op(e3, e3) = unit) | ~ sP43), inference(definition_unfolding, [], [f152, f128])).
fof(f152, plain, (~ (e0 = op(e3, e3)) | ~ sP43), inference(cnf_transformation, [], [f62])).
fof(f62, plain, ((~ (e0 = op(e3, e3)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e0, e3))) | ~ sP43), inference(nnf_transformation, [], [f50])).
fof(f50, plain, ((~ (e0 = op(e3, e3)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e0, e3))) | ~ sP43), inference(usedef, [], [e50])).
fof(e50, plain, (sP43 <=> (~ (e0 = op(e3, e3)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP43])])).
fof(f764, plain, (sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f763, f479])).
fof(f479, plain, ~ sP44, inference(subsumption_resolution, [], [f147, f126])).
fof(f147, plain, (~ (e1 = op(e3, e2)) | ~ sP44), inference(cnf_transformation, [], [f61])).
fof(f61, plain, ((~ (e1 = op(e3, e3)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e3, e0))) | ~ sP44), inference(nnf_transformation, [], [f51])).
fof(f51, plain, ((~ (e1 = op(e3, e3)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e3, e0))) | ~ sP44), inference(usedef, [], [e51])).
fof(e51, plain, (sP44 <=> (~ (e1 = op(e3, e3)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP44])])).
fof(f763, plain, (sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f762, f477])).
fof(f477, plain, ~ sP45, inference(subsumption_resolution, [], [f143, f123])).
fof(f143, plain, (~ (e1 = op(e2, e3)) | ~ sP45), inference(cnf_transformation, [], [f60])).
fof(f60, plain, ((~ (e1 = op(e3, e3)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e0, e3))) | ~ sP45), inference(nnf_transformation, [], [f52])).
fof(f52, plain, ((~ (e1 = op(e3, e3)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e0, e3))) | ~ sP45), inference(usedef, [], [e52])).
fof(e52, plain, (sP45 <=> (~ (e1 = op(e3, e3)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP45])])).
fof(f762, plain, (sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f761, f475])).
fof(f475, plain, ~ sP46, inference(subsumption_resolution, [], [f138, f125])).
fof(f138, plain, (~ (e2 = op(e3, e1)) | ~ sP46), inference(cnf_transformation, [], [f59])).
fof(f59, plain, ((~ (e2 = op(e3, e3)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e3, e0))) | ~ sP46), inference(nnf_transformation, [], [f53])).
fof(f53, plain, ((~ (e2 = op(e3, e3)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e3, e0))) | ~ sP46), inference(usedef, [], [e53])).
fof(e53, plain, (sP46 <=> (~ (e2 = op(e3, e3)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP46])])).
fof(f761, plain, (sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f760, f472])).
fof(f472, plain, ~ sP47, inference(subsumption_resolution, [], [f134, f119])).
fof(f134, plain, (~ (e2 = op(e1, e3)) | ~ sP47), inference(cnf_transformation, [], [f58])).
fof(f58, plain, ((~ (e2 = op(e3, e3)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e0, e3))) | ~ sP47), inference(nnf_transformation, [], [f54])).
fof(f54, plain, ((~ (e2 = op(e3, e3)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e0, e3))) | ~ sP47), inference(usedef, [], [e54])).
fof(e54, plain, (sP47 <=> (~ (e2 = op(e3, e3)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP47])])).
fof(f760, plain, (sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f759, f469])).
fof(f469, plain, ~ sP48, inference(subsumption_resolution, [], [f342, f333])).
fof(f342, plain, (~ (e3 = op(e3, unit)) | ~ sP48), inference(definition_unfolding, [], [f129, f128])).
fof(f129, plain, (~ (e3 = op(e3, e0)) | ~ sP48), inference(cnf_transformation, [], [f57])).
fof(f57, plain, ((~ (e3 = op(e3, e3)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e3, e0))) | ~ sP48), inference(nnf_transformation, [], [f55])).
fof(f55, plain, ((~ (e3 = op(e3, e3)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e3, e0))) | ~ sP48), inference(usedef, [], [e55])).
fof(e55, plain, (sP48 <=> (~ (e3 = op(e3, e3)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP48])])).
fof(f759, plain, (sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f461, f338])).
fof(f461, plain, (~ (e3 = op(unit, e3)) | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | ~ (e3 = op(e3, unit)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(duplicate_literal_removal, [], [f460])).
fof(f460, plain, (~ (e3 = op(unit, e3)) | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (unit = op(unit, unit)) | ~ (unit = op(unit, unit)) | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(definition_unfolding, [], [f325, f128, f128, f128, f128, f128])).
fof(f325, plain, (~ (e3 = op(e0, e3)) | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (e0 = op(e0, unit)) | ~ (e0 = op(unit, e0)) | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f56])).
fof(f56, plain, ((~ (e3 = op(e3, e3)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e0, e3))) | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (e0 = op(e0, unit)) | ~ (e0 = op(unit, e0)) | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(definition_folding, [], [f6, e55, e54, e53, e52, e51, e50, e49, e48, e47, e46, e45, e44, e43, e42, e41, e40, e39, e38, e37, e36, e35, e34, e33, e32, e31, e30, e29, e28, e27, e26, e25, e24, e23, e22, e21, e20, e19, e18, e17, e16, e15, e14, e13, e12, e11, e10, e9, e8, e7])).
fof(f6, plain, ((~ (e3 = op(e3, e3)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e0, e3))) | (~ (e3 = op(e3, e3)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e3, e0))) | (~ (e2 = op(e3, e3)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e0, e3))) | (~ (e2 = op(e3, e3)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e3, e0))) | (~ (e1 = op(e3, e3)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e0, e3))) | (~ (e1 = op(e3, e3)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e3, e0))) | (~ (e0 = op(e3, e3)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e0, e3))) | (~ (e0 = op(e3, e3)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e3, e0))) | (~ (e3 = op(e3, e2)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e0, e2))) | (~ (e3 = op(e2, e3)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e2, e0))) | (~ (e2 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e0, e2))) | (~ (e2 = op(e2, e3)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e2, e0))) | (~ (e1 = op(e3, e2)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e0, e2))) | (~ (e1 = op(e2, e3)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e2, e0))) | (~ (e0 = op(e3, e2)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e0, e2))) | (~ (e0 = op(e2, e3)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e2, e0))) | (~ (e3 = op(e3, e1)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e0, e1))) | (~ (e3 = op(e1, e3)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e1, e0))) | (~ (e2 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e0, e1))) | (~ (e2 = op(e1, e3)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e1, e0))) | (~ (e1 = op(e3, e1)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e0, e1))) | (~ (e1 = op(e1, e3)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e1, e0))) | (~ (e0 = op(e3, e1)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e0, e1))) | (~ (e0 = op(e1, e3)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e1, e0))) | (~ (e3 = op(e3, e0)) & ~ (e3 = op(e2, e0)) & ~ (e3 = op(e1, e0)) & ~ (e3 = op(e0, e0))) | (~ (e3 = op(e0, e3)) & ~ (e3 = op(e0, e2)) & ~ (e3 = op(e0, e1)) & ~ (e3 = op(e0, e0))) | (~ (e2 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e2 = op(e1, e0)) & ~ (e2 = op(e0, e0))) | (~ (e2 = op(e0, e3)) & ~ (e2 = op(e0, e2)) & ~ (e2 = op(e0, e1)) & ~ (e2 = op(e0, e0))) | (~ (e1 = op(e3, e0)) & ~ (e1 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e1 = op(e0, e0))) | (~ (e1 = op(e0, e3)) & ~ (e1 = op(e0, e2)) & ~ (e1 = op(e0, e1)) & ~ (e1 = op(e0, e0))) | (~ (e0 = op(e3, e0)) & ~ (e0 = op(e2, e0)) & ~ (e0 = op(e1, e0)) & ~ (e0 = op(e0, e0))) | (~ (e0 = op(e0, e3)) & ~ (e0 = op(e0, e2)) & ~ (e0 = op(e0, e1)) & ~ (e0 = op(e0, e0))) | (~ (e3 = unit) & ~ (e2 = unit) & ~ (e1 = unit) & ~ (e0 = unit)) | ~ (e3 = op(e3, unit)) | ~ (e3 = op(unit, e3)) | ~ (e2 = op(e2, unit)) | ~ (e2 = op(unit, e2)) | ~ (e1 = op(e1, unit)) | ~ (e1 = op(unit, e1)) | ~ (e0 = op(e0, unit)) | ~ (e0 = op(unit, e0)) | (~ (e3 = op(e3, e3)) & ~ (e2 = op(e3, e3)) & ~ (e1 = op(e3, e3)) & ~ (e0 = op(e3, e3))) | (~ (e3 = op(e3, e2)) & ~ (e2 = op(e3, e2)) & ~ (e1 = op(e3, e2)) & ~ (e0 = op(e3, e2))) | (~ (e3 = op(e3, e1)) & ~ (e2 = op(e3, e1)) & ~ (e1 = op(e3, e1)) & ~ (e0 = op(e3, e1))) | (~ (e3 = op(e3, e0)) & ~ (e2 = op(e3, e0)) & ~ (e1 = op(e3, e0)) & ~ (e0 = op(e3, e0))) | (~ (e3 = op(e2, e3)) & ~ (e2 = op(e2, e3)) & ~ (e1 = op(e2, e3)) & ~ (e0 = op(e2, e3))) | (~ (e3 = op(e2, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e2, e2)) & ~ (e0 = op(e2, e2))) | (~ (e3 = op(e2, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e2, e1)) & ~ (e0 = op(e2, e1))) | (~ (e3 = op(e2, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e2, e0)) & ~ (e0 = op(e2, e0))) | (~ (e3 = op(e1, e3)) & ~ (e2 = op(e1, e3)) & ~ (e1 = op(e1, e3)) & ~ (e0 = op(e1, e3))) | (~ (e3 = op(e1, e2)) & ~ (e2 = op(e1, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e1, e2))) | (~ (e3 = op(e1, e1)) & ~ (e2 = op(e1, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e1, e1))) | (~ (e3 = op(e1, e0)) & ~ (e2 = op(e1, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e1, e0))) | (~ (e3 = op(e0, e3)) & ~ (e2 = op(e0, e3)) & ~ (e1 = op(e0, e3)) & ~ (e0 = op(e0, e3))) | (~ (e3 = op(e0, e2)) & ~ (e2 = op(e0, e2)) & ~ (e1 = op(e0, e2)) & ~ (e0 = op(e0, e2))) | (~ (e3 = op(e0, e1)) & ~ (e2 = op(e0, e1)) & ~ (e1 = op(e0, e1)) & ~ (e0 = op(e0, e1))) | (~ (e3 = op(e0, e0)) & ~ (e2 = op(e0, e0)) & ~ (e1 = op(e0, e0)) & ~ (e0 = op(e0, e0))) | ((~ (e3 = op(e3, e3)) | ~ (e3 = op(e2, e2)) | ~ (e3 = op(e1, e1)) | ~ (e3 = op(e0, e0))) & (~ (e2 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e2 = op(e1, e1)) | ~ (e2 = op(e0, e0))) & (~ (e1 = op(e3, e3)) | ~ (e1 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e1 = op(e0, e0))) & (~ (e0 = op(e3, e3)) | ~ (e0 = op(e2, e2)) | ~ (e0 = op(e1, e1)) | ~ (e0 = op(e0, e0))))), inference(ennf_transformation, [], [f5])).
fof(f5, plain, ~ (((e3 = op(e3, e3)) | (e3 = op(e2, e3)) | (e3 = op(e1, e3)) | (e3 = op(e0, e3))) & ((e3 = op(e3, e3)) | (e3 = op(e3, e2)) | (e3 = op(e3, e1)) | (e3 = op(e3, e0))) & ((e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(e0, e3))) & ((e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, e0))) & ((e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(e0, e3))) & ((e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, e0))) & ((e0 = op(e3, e3)) | (e0 = op(e2, e3)) | (e0 = op(e1, e3)) | (e0 = op(e0, e3))) & ((e0 = op(e3, e3)) | (e0 = op(e3, e2)) | (e0 = op(e3, e1)) | (e0 = op(e3, e0))) & ((e3 = op(e3, e2)) | (e3 = op(e2, e2)) | (e3 = op(e1, e2)) | (e3 = op(e0, e2))) & ((e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, e0))) & ((e2 = op(e3, e2)) | (e2 = op(e2, e2)) | (e2 = op(e1, e2)) | (e2 = op(e0, e2))) & ((e2 = op(e2, e3)) | (e2 = op(e2, e2)) | (e2 = op(e2, e1)) | (e2 = op(e2, e0))) & ((e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(e0, e2))) & ((e1 = op(e2, e3)) | (e1 = op(e2, e2)) | (e1 = op(e2, e1)) | (e1 = op(e2, e0))) & ((e0 = op(e3, e2)) | (e0 = op(e2, e2)) | (e0 = op(e1, e2)) | (e0 = op(e0, e2))) & ((e0 = op(e2, e3)) | (e0 = op(e2, e2)) | (e0 = op(e2, e1)) | (e0 = op(e2, e0))) & ((e3 = op(e3, e1)) | (e3 = op(e2, e1)) | (e3 = op(e1, e1)) | (e3 = op(e0, e1))) & ((e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, e0))) & ((e2 = op(e3, e1)) | (e2 = op(e2, e1)) | (e2 = op(e1, e1)) | (e2 = op(e0, e1))) & ((e2 = op(e1, e3)) | (e2 = op(e1, e2)) | (e2 = op(e1, e1)) | (e2 = op(e1, e0))) & ((e1 = op(e3, e1)) | (e1 = op(e2, e1)) | (e1 = op(e1, e1)) | (e1 = op(e0, e1))) & ((e1 = op(e1, e3)) | (e1 = op(e1, e2)) | (e1 = op(e1, e1)) | (e1 = op(e1, e0))) & ((e0 = op(e3, e1)) | (e0 = op(e2, e1)) | (e0 = op(e1, e1)) | (e0 = op(e0, e1))) & ((e0 = op(e1, e3)) | (e0 = op(e1, e2)) | (e0 = op(e1, e1)) | (e0 = op(e1, e0))) & ((e3 = op(e3, e0)) | (e3 = op(e2, e0)) | (e3 = op(e1, e0)) | (e3 = op(e0, e0))) & ((e3 = op(e0, e3)) | (e3 = op(e0, e2)) | (e3 = op(e0, e1)) | (e3 = op(e0, e0))) & ((e2 = op(e3, e0)) | (e2 = op(e2, e0)) | (e2 = op(e1, e0)) | (e2 = op(e0, e0))) & ((e2 = op(e0, e3)) | (e2 = op(e0, e2)) | (e2 = op(e0, e1)) | (e2 = op(e0, e0))) & ((e1 = op(e3, e0)) | (e1 = op(e2, e0)) | (e1 = op(e1, e0)) | (e1 = op(e0, e0))) & ((e1 = op(e0, e3)) | (e1 = op(e0, e2)) | (e1 = op(e0, e1)) | (e1 = op(e0, e0))) & ((e0 = op(e3, e0)) | (e0 = op(e2, e0)) | (e0 = op(e1, e0)) | (e0 = op(e0, e0))) & ((e0 = op(e0, e3)) | (e0 = op(e0, e2)) | (e0 = op(e0, e1)) | (e0 = op(e0, e0))) & ((e3 = unit) | (e2 = unit) | (e1 = unit) | (e0 = unit)) & (e3 = op(e3, unit)) & (e3 = op(unit, e3)) & (e2 = op(e2, unit)) & (e2 = op(unit, e2)) & (e1 = op(e1, unit)) & (e1 = op(unit, e1)) & (e0 = op(e0, unit)) & (e0 = op(unit, e0)) & ((e3 = op(e3, e3)) | (e2 = op(e3, e3)) | (e1 = op(e3, e3)) | (e0 = op(e3, e3))) & ((e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (e0 = op(e3, e2))) & ((e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (e0 = op(e3, e1))) & ((e3 = op(e3, e0)) | (e2 = op(e3, e0)) | (e1 = op(e3, e0)) | (e0 = op(e3, e0))) & ((e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))) & ((e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (e0 = op(e2, e2))) & ((e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))) & ((e3 = op(e2, e0)) | (e2 = op(e2, e0)) | (e1 = op(e2, e0)) | (e0 = op(e2, e0))) & ((e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (e0 = op(e1, e3))) & ((e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (e0 = op(e1, e2))) & ((e3 = op(e1, e1)) | (e2 = op(e1, e1)) | (e1 = op(e1, e1)) | (e0 = op(e1, e1))) & ((e3 = op(e1, e0)) | (e2 = op(e1, e0)) | (e1 = op(e1, e0)) | (e0 = op(e1, e0))) & ((e3 = op(e0, e3)) | (e2 = op(e0, e3)) | (e1 = op(e0, e3)) | (e0 = op(e0, e3))) & ((e3 = op(e0, e2)) | (e2 = op(e0, e2)) | (e1 = op(e0, e2)) | (e0 = op(e0, e2))) & ((e3 = op(e0, e1)) | (e2 = op(e0, e1)) | (e1 = op(e0, e1)) | (e0 = op(e0, e1))) & ((e3 = op(e0, e0)) | (e2 = op(e0, e0)) | (e1 = op(e0, e0)) | (e0 = op(e0, e0))) & (((e3 = op(e3, e3)) & (e3 = op(e2, e2)) & (e3 = op(e1, e1)) & (e3 = op(e0, e0))) | ((e2 = op(e3, e3)) & (e2 = op(e2, e2)) & (e2 = op(e1, e1)) & (e2 = op(e0, e0))) | ((e1 = op(e3, e3)) & (e1 = op(e2, e2)) & (e1 = op(e1, e1)) & (e1 = op(e0, e0))) | ((e0 = op(e3, e3)) & (e0 = op(e2, e2)) & (e0 = op(e1, e1)) & (e0 = op(e0, e0))))), inference(negated_conjecture, [], [f4])).
fof(f4, plain, ~ (((e3 = op(e3, e3)) | (e3 = op(e2, e3)) | (e3 = op(e1, e3)) | (e3 = op(e0, e3))) & ((e3 = op(e3, e3)) | (e3 = op(e3, e2)) | (e3 = op(e3, e1)) | (e3 = op(e3, e0))) & ((e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(e0, e3))) & ((e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, e0))) & ((e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(e0, e3))) & ((e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, e0))) & ((e0 = op(e3, e3)) | (e0 = op(e2, e3)) | (e0 = op(e1, e3)) | (e0 = op(e0, e3))) & ((e0 = op(e3, e3)) | (e0 = op(e3, e2)) | (e0 = op(e3, e1)) | (e0 = op(e3, e0))) & ((e3 = op(e3, e2)) | (e3 = op(e2, e2)) | (e3 = op(e1, e2)) | (e3 = op(e0, e2))) & ((e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, e0))) & ((e2 = op(e3, e2)) | (e2 = op(e2, e2)) | (e2 = op(e1, e2)) | (e2 = op(e0, e2))) & ((e2 = op(e2, e3)) | (e2 = op(e2, e2)) | (e2 = op(e2, e1)) | (e2 = op(e2, e0))) & ((e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(e0, e2))) & ((e1 = op(e2, e3)) | (e1 = op(e2, e2)) | (e1 = op(e2, e1)) | (e1 = op(e2, e0))) & ((e0 = op(e3, e2)) | (e0 = op(e2, e2)) | (e0 = op(e1, e2)) | (e0 = op(e0, e2))) & ((e0 = op(e2, e3)) | (e0 = op(e2, e2)) | (e0 = op(e2, e1)) | (e0 = op(e2, e0))) & ((e3 = op(e3, e1)) | (e3 = op(e2, e1)) | (e3 = op(e1, e1)) | (e3 = op(e0, e1))) & ((e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, e0))) & ((e2 = op(e3, e1)) | (e2 = op(e2, e1)) | (e2 = op(e1, e1)) | (e2 = op(e0, e1))) & ((e2 = op(e1, e3)) | (e2 = op(e1, e2)) | (e2 = op(e1, e1)) | (e2 = op(e1, e0))) & ((e1 = op(e3, e1)) | (e1 = op(e2, e1)) | (e1 = op(e1, e1)) | (e1 = op(e0, e1))) & ((e1 = op(e1, e3)) | (e1 = op(e1, e2)) | (e1 = op(e1, e1)) | (e1 = op(e1, e0))) & ((e0 = op(e3, e1)) | (e0 = op(e2, e1)) | (e0 = op(e1, e1)) | (e0 = op(e0, e1))) & ((e0 = op(e1, e3)) | (e0 = op(e1, e2)) | (e0 = op(e1, e1)) | (e0 = op(e1, e0))) & ((e3 = op(e3, e0)) | (e3 = op(e2, e0)) | (e3 = op(e1, e0)) | (e3 = op(e0, e0))) & ((e3 = op(e0, e3)) | (e3 = op(e0, e2)) | (e3 = op(e0, e1)) | (e3 = op(e0, e0))) & ((e2 = op(e3, e0)) | (e2 = op(e2, e0)) | (e2 = op(e1, e0)) | (e2 = op(e0, e0))) & ((e2 = op(e0, e3)) | (e2 = op(e0, e2)) | (e2 = op(e0, e1)) | (e2 = op(e0, e0))) & ((e1 = op(e3, e0)) | (e1 = op(e2, e0)) | (e1 = op(e1, e0)) | (e1 = op(e0, e0))) & ((e1 = op(e0, e3)) | (e1 = op(e0, e2)) | (e1 = op(e0, e1)) | (e1 = op(e0, e0))) & ((e0 = op(e3, e0)) | (e0 = op(e2, e0)) | (e0 = op(e1, e0)) | (e0 = op(e0, e0))) & ((e0 = op(e0, e3)) | (e0 = op(e0, e2)) | (e0 = op(e0, e1)) | (e0 = op(e0, e0))) & ((e3 = unit) | (e2 = unit) | (e1 = unit) | (e0 = unit)) & (e3 = op(e3, unit)) & (e3 = op(unit, e3)) & (e2 = op(e2, unit)) & (e2 = op(unit, e2)) & (e1 = op(e1, unit)) & (e1 = op(unit, e1)) & (e0 = op(e0, unit)) & (e0 = op(unit, e0)) & ((e3 = op(e3, e3)) | (e2 = op(e3, e3)) | (e1 = op(e3, e3)) | (e0 = op(e3, e3))) & ((e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (e0 = op(e3, e2))) & ((e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (e0 = op(e3, e1))) & ((e3 = op(e3, e0)) | (e2 = op(e3, e0)) | (e1 = op(e3, e0)) | (e0 = op(e3, e0))) & ((e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))) & ((e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (e0 = op(e2, e2))) & ((e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))) & ((e3 = op(e2, e0)) | (e2 = op(e2, e0)) | (e1 = op(e2, e0)) | (e0 = op(e2, e0))) & ((e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (e0 = op(e1, e3))) & ((e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (e0 = op(e1, e2))) & ((e3 = op(e1, e1)) | (e2 = op(e1, e1)) | (e1 = op(e1, e1)) | (e0 = op(e1, e1))) & ((e3 = op(e1, e0)) | (e2 = op(e1, e0)) | (e1 = op(e1, e0)) | (e0 = op(e1, e0))) & ((e3 = op(e0, e3)) | (e2 = op(e0, e3)) | (e1 = op(e0, e3)) | (e0 = op(e0, e3))) & ((e3 = op(e0, e2)) | (e2 = op(e0, e2)) | (e1 = op(e0, e2)) | (e0 = op(e0, e2))) & ((e3 = op(e0, e1)) | (e2 = op(e0, e1)) | (e1 = op(e0, e1)) | (e0 = op(e0, e1))) & ((e3 = op(e0, e0)) | (e2 = op(e0, e0)) | (e1 = op(e0, e0)) | (e0 = op(e0, e0))) & (((e3 = op(e3, e3)) & (e3 = op(e2, e2)) & (e3 = op(e1, e1)) & (e3 = op(e0, e0))) | ((e2 = op(e3, e3)) & (e2 = op(e2, e2)) & (e2 = op(e1, e1)) & (e2 = op(e0, e0))) | ((e1 = op(e3, e3)) & (e1 = op(e2, e2)) & (e1 = op(e1, e1)) & (e1 = op(e0, e0))) | ((e0 = op(e3, e3)) & (e0 = op(e2, e2)) & (e0 = op(e1, e1)) & (e0 = op(e0, e0))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG043+1.p', co1)).