fof(f8830, plain, $false, inference(avatar_sat_refutation, [], [f385, f402, f419, f436, f453, f470, f487, f504, f521, f538, f555, f572, f589, f606, f623, f640, f641, f642, f643, f644, f645, f646, f647, f648, f649, f650, f651, f652, f653, f654, f655, f656, f657, f658, f659, f660, f661, f662, f663, f664, f665, f666, f667, f668, f669, f670, f671, f672, f681, f686, f691, f696, f701, f715, f720, f725, f730, f739, f749, f754, f759, f768, f773, f778, f783, f784, f793, f798, f803, f808, f817, f822, f832, f837, f846, f851, f856, f866, f880, f885, f890, f891, f900, f905, f910, f915, f916, f925, f930, f935, f940, f949, f954, f959, f969, f978, f983, f988, f993, f994, f1003, f1008, f1013, f1018, f1019, f1028, f1033, f1038, f1044, f1053, f1058, f1063, f1068, f1077, f1082, f1087, f1096, f1101, f1106, f1115, f1120, f1125, f1134, f1139, f1144, f1153, f1154, f1164, f1169, f1174, f1183, f1188, f1193, f1202, f1207, f1212, f1221, f1226, f1231, f1240, f1241, f1251, f1256, f1261, f1270, f1275, f1280, f1289, f1294, f1299, f1308, f1313, f1318, f1327, f1328, f1334, f1339, f1349, f1354, f1355, f1356, f1365, f1370, f1375, f1380, f1385, f1386, f1395, f1400, f1405, f1410, f1415, f1416, f1425, f1430, f1435, f1440, f1445, f1446, f1455, f1460, f1465, f1470, f1475, f1476, f1501, f1521, f1524, f1531, f1565, f1568, f1581, f1598, f1601, f1604, f1613, f1616, f1623, f1626, f1633, f1638, f1641, f1644, f1661, f1672, f1682, f1685, f1696, f1713, f1715, f1718, f1723, f1732, f1774, f1781, f1821, f1832, f1838, f1842, f1845, f1853, f1859, f1908, f1916, f1925, f1930, f1945, f1955, f1977, f1979, f1982, f2053, f2062, f2065, f2068, f2074, f2075, f2093, f2112, f2130, f2151, f2165, f2169, f2181, f2192, f2207, f2237, f2251, f2258, f2268, f2270, f2272, f2289, f2292, f2297, f2322, f2327, f2355, f2372, f2420, f2432, f2445, f2457, f2483, f2485, f2503, f2521, f2525, f2534, f2537, f2550, f2553, f2558, f2563, f2569, f2570, f2572, f2573, f2590, f2592, f2605, f2617, f2636, f2639, f2663, f2673, f2676, f2694, f2730, f2756, f2763, f2776, f2797, f2801, f2888, f2920, f2943, f2948, f2949, f2959, f2962, f2980, f3001, f3021, f3038, f3053, f3098, f3101, f3128, f3142, f3146, f3168, f3172, f3178, f3186, f3197, f3209, f3216, f3243, f3245, f3247, f3296, f3312, f3316, f3336, f3338, f3373, f3378, f3410, f3420, f3454, f3460, f3473, f3502, f3507, f3511, f3523, f3538, f3544, f3550, f3575, f3576, f3588, f3590, f3604, f3607, f3631, f3681, f3692, f3700, f3710, f3712, f3733, f3743, f3750, f3784, f3796, f3824, f3832, f3849, f3876, f3878, f3902, f3905, f3918, f3921, f3927, f3961, f4035, f4036, f4038, f4080, f4083, f4085, f4091, f4095, f4105, f4111, f4121, f4163, f4175, f4217, f4229, f4242, f4245, f4247, f4264, f4305, f4323, f4338, f4342, f4346, f4357, f4359, f4363, f4370, f4388, f4390, f4408, f4429, f4446, f4448, f4529, f4544, f4550, f4551, f4553, f4557, f4559, f4572, f4589, f4591, f4609, f4634, f4642, f4656, f4677, f4678, f4680, f4681, f4695, f4723, f4726, f4729, f4750, f4779, f4781, f4792, f4810, f4819, f4867, f4868, f4870, f4881, f4885, f4893, f4895, f4902, f4912, f4914, f4916, f4962, f4964, f4979, f4980, f5026, f5060, f5070, f5071, f5089, f5101, f5106, f5107, f5109, f5147, f5151, f5154, f5163, f5173, f5226, f5229, f5230, f5232, f5240, f5242, f5259, f5274, f5278, f5296, f5302, f5315, f5334, f5337, f5342, f5388, f5393, f5396, f5440, f5452, f5454, f5458, f5461, f5531, f5532, f5534, f5547, f5549, f5553, f5565, f5568, f5569, f5570, f5585, f5606, f5611, f5619, f5621, f5690, f5693, f5697, f5727, f5729, f5754, f5783, f5803, f5823, f5833, f5834, f5836, f5850, f5860, f5865, f5875, f5880, f5907, f5908, f5917, f5920, f5922, f5937, f5938, f5940, f5941, f5943, f5954, f5972, f5992, f6008, f6061, f6064, f6065, f6107, f6109, f6123, f6125, f6140, f6156, f6167, f6187, f6202, f6227, f6233, f6248, f6268, f6286, f6313, f6314, f6363, f6370, f6391, f6403, f6407, f6431, f6443, f6464, f6468, f6497, f6508, f6511, f6524, f6527, f6529, f6530, f6535, f6539, f6541, f6548, f6572, f6577, f6579, f6582, f6602, f6629, f6630, f6658, f6660, f6662, f6674, f6698, f6730, f6732, f6750, f6780, f6781, f6785, f6817, f6831, f6845, f6846, f6879, f6880, f6882, f6896, f6910, f6935, f6940, f6947, f6963, f6994, f7011, f7016, f7027, f7030, f7042, f7045, f7065, f7075, f7115, f7116, f7117, f7120, f7136, f7137, f7138, f7153, f7154, f7166, f7167, f7171, f7187, f7188, f7190, f7191, f7203, f7217, f7242, f7246, f7247, f7273, f7276, f7285, f7292, f7304, f7307, f7318, f7321, f7350, f7351, f7355, f7369, f7375, f7424, f7438, f7446, f7453, f7455, f7475, f7479, f7481, f7485, f7496, f7509, f7511, f7540, f7541, f7556, f7557, f7570, f7584, f7601, f7610, f7629, f7640, f7655, f7687, f7701, f7705, f7724, f7738, f7742, f7752, f7753, f7768, f7769, f7770, f7773, f7774, f7785, f7786, f7798, f7803, f7813, f7814, f7831, f7858, f7861, f7865, f7867, f7871, f7873, f7891, f7894, f7901, f7907, f7909, f7912, f7926, f7927, f7937, f7939, f7945, f7957, f7960, f7962, f7967, f7969, f7972, f7974, f7975, f7976, f7979, f7980, f7982, f7984, f7986, f7987, f7996, f7998, f8000, f8007, f8015, f8018, f8020, f8033, f8076, f8090, f8096, f8099, f8101, f8103, f8105, f8112, f8125, f8138, f8152, f8153, f8187, f8188, f8189, f8204, f8206, f8219, f8220, f8233, f8256, f8258, f8283, f8289, f8296, f8299, f8322, f8323, f8324, f8326, f8337, f8338, f8340, f8351, f8352, f8355, f8371, f8372, f8387, f8417, f8418, f8421, f8430, f8445, f8446, f8473, f8479, f8488, f8491, f8498, f8500, f8504, f8511, f8537, f8538, f8553, f8554, f8568, f8569, f8570, f8589, f8590, f8609, f8610, f8611, f8613, f8628, f8629, f8638, f8648, f8665, f8668, f8673, f8681, f8692, f8694, f8700, f8702, f8711, f8715, f8726, f8732, f8743, f8744, f8746, f8747, f8761, f8762, f8763, f8777, f8796, f8797, f8810])).
fof(f8810, plain, (~ spl30_7 | ~ spl30_15), inference(avatar_split_clause, [], [f8802, f429, f395])).
fof(f395, plain, (spl30_7 <=> (e2 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_7])])).
fof(f429, plain, (spl30_15 <=> (e2 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_15])])).
fof(f8802, plain, (~ (e2 = op(e3, e2)) | ~ spl30_15), inference(backward_demodulation, [], [f206, f431])).
fof(f431, plain, ((e2 = op(e3, e0)) | ~ spl30_15), inference(avatar_component_clause, [], [f429])).
fof(f206, plain, ~ (op(e3, e0) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f3, plain, (~ (op(e3, e2) = op(e3, e3)) & ~ (op(e3, e1) = op(e3, e3)) & ~ (op(e3, e1) = op(e3, e2)) & ~ (op(e3, e0) = op(e3, e3)) & ~ (op(e3, e0) = op(e3, e2)) & ~ (op(e3, e0) = op(e3, e1)) & ~ (op(e2, e2) = op(e2, e3)) & ~ (op(e2, e1) = op(e2, e3)) & ~ (op(e2, e1) = op(e2, e2)) & ~ (op(e2, e0) = op(e2, e3)) & ~ (op(e2, e0) = op(e2, e2)) & ~ (op(e2, e0) = op(e2, e1)) & ~ (op(e1, e2) = op(e1, e3)) & ~ (op(e1, e1) = op(e1, e3)) & ~ (op(e1, e1) = op(e1, e2)) & ~ (op(e1, e0) = op(e1, e3)) & ~ (op(e1, e0) = op(e1, e2)) & ~ (op(e1, e0) = op(e1, e1)) & ~ (op(e0, e2) = op(e0, e3)) & ~ (op(e0, e1) = op(e0, e3)) & ~ (op(e0, e1) = op(e0, e2)) & ~ (op(e0, e0) = op(e0, e3)) & ~ (op(e0, e0) = op(e0, e2)) & ~ (op(e0, e0) = op(e0, e1)) & ~ (op(e2, e3) = op(e3, e3)) & ~ (op(e1, e3) = op(e3, e3)) & ~ (op(e1, e3) = op(e2, e3)) & ~ (op(e0, e3) = op(e3, e3)) & ~ (op(e0, e3) = op(e2, e3)) & ~ (op(e0, e3) = op(e1, e3)) & ~ (op(e2, e2) = op(e3, e2)) & ~ (op(e1, e2) = op(e3, e2)) & ~ (op(e1, e2) = op(e2, e2)) & ~ (op(e0, e2) = op(e3, e2)) & ~ (op(e0, e2) = op(e2, e2)) & ~ (op(e0, e2) = op(e1, e2)) & ~ (op(e2, e1) = op(e3, e1)) & ~ (op(e1, e1) = op(e3, e1)) & ~ (op(e1, e1) = op(e2, e1)) & ~ (op(e0, e1) = op(e3, e1)) & ~ (op(e0, e1) = op(e2, e1)) & ~ (op(e0, e1) = op(e1, e1)) & ~ (op(e2, e0) = op(e3, e0)) & ~ (op(e1, e0) = op(e3, e0)) & ~ (op(e1, e0) = op(e2, e0)) & ~ (op(e0, e0) = op(e3, e0)) & ~ (op(e0, e0) = op(e2, e0)) & ~ (op(e0, e0) = op(e1, e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG157+1.p', ax3)).
fof(f8797, plain, (spl30_10 | ~ spl30_34 | ~ spl30_76), inference(avatar_split_clause, [], [f8786, f727, f510, f408])).
fof(f408, plain, (spl30_10 <=> (e1 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_10])])).
fof(f510, plain, (spl30_34 <=> (e1 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_34])])).
fof(f727, plain, (spl30_76 <=> (op(e1, e3) = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_76])])).
fof(f8786, plain, ((e1 = op(e3, e1)) | (~ spl30_34 | ~ spl30_76)), inference(backward_demodulation, [], [f729, f512])).
fof(f512, plain, ((e1 = op(e1, e3)) | ~ spl30_34), inference(avatar_component_clause, [], [f510])).
fof(f729, plain, ((op(e1, e3) = op(e3, e1)) | ~ spl30_76), inference(avatar_component_clause, [], [f727])).
fof(f8796, plain, (~ spl30_2 | ~ spl30_34), inference(avatar_split_clause, [], [f8783, f510, f374])).
fof(f374, plain, (spl30_2 <=> (e1 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_2])])).
fof(f8783, plain, (~ (e1 = op(e3, e3)) | ~ spl30_34), inference(backward_demodulation, [], [f185, f512])).
fof(f185, plain, ~ (op(e1, e3) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f8777, plain, (~ spl30_7 | ~ spl30_39), inference(avatar_split_clause, [], [f8767, f531, f395])).
fof(f531, plain, (spl30_39 <=> (e2 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_39])])).
fof(f8767, plain, (~ (e2 = op(e3, e2)) | ~ spl30_39), inference(backward_demodulation, [], [f179, f533])).
fof(f533, plain, ((e2 = op(e1, e2)) | ~ spl30_39), inference(avatar_component_clause, [], [f531])).
fof(f179, plain, ~ (op(e1, e2) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f8763, plain, (~ spl30_16 | ~ spl30_41 | spl30_116), inference(avatar_split_clause, [], [f8754, f922, f540, f433])).
fof(f433, plain, (spl30_16 <=> (e3 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_16])])).
fof(f540, plain, (spl30_41 <=> (e0 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_41])])).
fof(f922, plain, (spl30_116 <=> (e3 = op(e3, op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl30_116])])).
fof(f8754, plain, (~ (e3 = op(e3, e0)) | (~ spl30_41 | spl30_116)), inference(backward_demodulation, [], [f924, f542])).
fof(f542, plain, ((e0 = op(e1, e1)) | ~ spl30_41), inference(avatar_component_clause, [], [f540])).
fof(f924, plain, (~ (e3 = op(e3, op(e1, e1))) | spl30_116), inference(avatar_component_clause, [], [f922])).
fof(f8762, plain, (~ spl30_33 | ~ spl30_41), inference(avatar_split_clause, [], [f8750, f540, f506])).
fof(f506, plain, (spl30_33 <=> (e0 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_33])])).
fof(f8750, plain, (~ (e0 = op(e1, e3)) | ~ spl30_41), inference(backward_demodulation, [], [f197, f542])).
fof(f197, plain, ~ (op(e1, e1) = op(e1, e3)), inference(cnf_transformation, [], [f3])).
fof(f8761, plain, (~ spl30_9 | ~ spl30_41), inference(avatar_split_clause, [], [f8748, f540, f404])).
fof(f404, plain, (spl30_9 <=> (e0 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_9])])).
fof(f8748, plain, (~ (e0 = op(e3, e1)) | ~ spl30_41), inference(backward_demodulation, [], [f173, f542])).
fof(f173, plain, ~ (op(e1, e1) = op(e3, e1)), inference(cnf_transformation, [], [f3])).
fof(f8747, plain, (~ spl30_33 | ~ spl30_48 | spl30_185), inference(avatar_split_clause, [], [f8742, f1258, f569, f506])).
fof(f569, plain, (spl30_48 <=> (e3 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_48])])).
fof(f1258, plain, (spl30_185 <=> (e0 = op(e1, op(e1, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl30_185])])).
fof(f8742, plain, (~ (e0 = op(e1, e3)) | (~ spl30_48 | spl30_185)), inference(backward_demodulation, [], [f1260, f571])).
fof(f571, plain, ((e3 = op(e1, e0)) | ~ spl30_48), inference(avatar_component_clause, [], [f569])).
fof(f1260, plain, (~ (e0 = op(e1, op(e1, e0))) | spl30_185), inference(avatar_component_clause, [], [f1258])).
fof(f8746, plain, (~ spl30_9 | ~ spl30_48 | spl30_184), inference(avatar_split_clause, [], [f8741, f1253, f569, f404])).
fof(f1253, plain, (spl30_184 <=> (e0 = op(op(e1, e0), e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_184])])).
fof(f8741, plain, (~ (e0 = op(e3, e1)) | (~ spl30_48 | spl30_184)), inference(backward_demodulation, [], [f1254, f571])).
fof(f1254, plain, (~ (e0 = op(op(e1, e0), e1)) | spl30_184), inference(avatar_component_clause, [], [f1253])).
fof(f8744, plain, (~ spl30_40 | ~ spl30_48), inference(avatar_split_clause, [], [f8735, f569, f535])).
fof(f535, plain, (spl30_40 <=> (e3 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_40])])).
fof(f8735, plain, (~ (e3 = op(e1, e2)) | ~ spl30_48), inference(backward_demodulation, [], [f194, f571])).
fof(f194, plain, ~ (op(e1, e0) = op(e1, e2)), inference(cnf_transformation, [], [f3])).
fof(f8743, plain, (~ spl30_16 | ~ spl30_48), inference(avatar_split_clause, [], [f8733, f569, f433])).
fof(f8733, plain, (~ (e3 = op(e3, e0)) | ~ spl30_48), inference(backward_demodulation, [], [f167, f571])).
fof(f167, plain, ~ (op(e1, e0) = op(e3, e0)), inference(cnf_transformation, [], [f3])).
fof(f8732, plain, (spl30_8 | ~ spl30_20 | ~ spl30_70), inference(avatar_split_clause, [], [f8731, f698, f450, f399])).
fof(f399, plain, (spl30_8 <=> (e3 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_8])])).
fof(f450, plain, (spl30_20 <=> (e3 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_20])])).
fof(f698, plain, (spl30_70 <=> (op(e2, e3) = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_70])])).
fof(f8731, plain, ((e3 = op(e3, e2)) | (~ spl30_20 | ~ spl30_70)), inference(forward_demodulation, [], [f700, f452])).
fof(f452, plain, ((e3 = op(e2, e3)) | ~ spl30_20), inference(avatar_component_clause, [], [f450])).
fof(f700, plain, ((op(e2, e3) = op(e3, e2)) | ~ spl30_70), inference(avatar_component_clause, [], [f698])).
fof(f8726, plain, (spl30_15 | ~ spl30_51 | ~ spl30_82), inference(avatar_split_clause, [], [f8725, f756, f582, f429])).
fof(f582, plain, (spl30_51 <=> (e2 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_51])])).
fof(f756, plain, (spl30_82 <=> (op(e0, e3) = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_82])])).
fof(f8725, plain, ((e2 = op(e3, e0)) | (~ spl30_51 | ~ spl30_82)), inference(forward_demodulation, [], [f758, f584])).
fof(f584, plain, ((e2 = op(e0, e3)) | ~ spl30_51), inference(avatar_component_clause, [], [f582])).
fof(f758, plain, ((op(e0, e3) = op(e3, e0)) | ~ spl30_82), inference(avatar_component_clause, [], [f756])).
fof(f8715, plain, (~ spl30_42 | ~ spl30_54 | spl30_134), inference(avatar_split_clause, [], [f8714, f1010, f595, f544])).
fof(f544, plain, (spl30_42 <=> (e1 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_42])])).
fof(f595, plain, (spl30_54 <=> (e1 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_54])])).
fof(f1010, plain, (spl30_134 <=> (e1 = op(e1, op(e0, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl30_134])])).
fof(f8714, plain, (~ (e1 = op(e1, e1)) | (~ spl30_54 | spl30_134)), inference(forward_demodulation, [], [f1012, f597])).
fof(f597, plain, ((e1 = op(e0, e2)) | ~ spl30_54), inference(avatar_component_clause, [], [f595])).
fof(f1012, plain, (~ (e1 = op(e1, op(e0, e2))) | spl30_134), inference(avatar_component_clause, [], [f1010])).
fof(f8711, plain, (~ spl30_16 | ~ spl30_61 | spl30_142), inference(avatar_split_clause, [], [f8710, f1050, f625, f433])).
fof(f625, plain, (spl30_61 <=> (e0 = op(e0, e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_61])])).
fof(f1050, plain, (spl30_142 <=> (e3 = op(e3, op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl30_142])])).
fof(f8710, plain, (~ (e3 = op(e3, e0)) | (~ spl30_61 | spl30_142)), inference(forward_demodulation, [], [f1052, f627])).
fof(f627, plain, ((e0 = op(e0, e0)) | ~ spl30_61), inference(avatar_component_clause, [], [f625])).
fof(f1052, plain, (~ (e3 = op(e3, op(e0, e0))) | spl30_142), inference(avatar_component_clause, [], [f1050])).
fof(f8702, plain, (~ spl30_47 | ~ spl30_30 | spl30_169), inference(avatar_split_clause, [], [f8701, f1180, f493, f565])).
fof(f565, plain, (spl30_47 <=> (e2 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_47])])).
fof(f493, plain, (spl30_30 <=> (e1 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_30])])).
fof(f1180, plain, (spl30_169 <=> (e2 = op(op(e2, e0), e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_169])])).
fof(f8701, plain, (~ (e2 = op(e1, e0)) | (~ spl30_30 | spl30_169)), inference(forward_demodulation, [], [f1182, f495])).
fof(f495, plain, ((e1 = op(e2, e0)) | ~ spl30_30), inference(avatar_component_clause, [], [f493])).
fof(f1182, plain, (~ (e2 = op(op(e2, e0), e0)) | spl30_169), inference(avatar_component_clause, [], [f1180])).
fof(f8700, plain, (~ spl30_30 | ~ spl30_51 | ~ spl30_188), inference(avatar_contradiction_clause, [], [f8699])).
fof(f8699, plain, ($false | (~ spl30_30 | ~ spl30_51 | ~ spl30_188)), inference(subsumption_resolution, [], [f8698, f215])).
fof(f215, plain, ~ (e1 = e3), inference(cnf_transformation, [], [f4])).
fof(f4, plain, (~ (e2 = e3) & ~ (e1 = e3) & ~ (e1 = e2) & ~ (e0 = e3) & ~ (e0 = e2) & ~ (e0 = e1)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG157+1.p', ax4)).
fof(f8698, plain, ((e1 = e3) | (~ spl30_30 | ~ spl30_51 | ~ spl30_188)), inference(forward_demodulation, [], [f8697, f495])).
fof(f8697, plain, ((e3 = op(e2, e0)) | (~ spl30_51 | ~ spl30_188)), inference(forward_demodulation, [], [f1274, f584])).
fof(f1274, plain, ((e3 = op(op(e0, e3), e0)) | ~ spl30_188), inference(avatar_component_clause, [], [f1272])).
fof(f1272, plain, (spl30_188 <=> (e3 = op(op(e0, e3), e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_188])])).
fof(f8694, plain, (~ spl30_47 | ~ spl30_54 | spl30_192), inference(avatar_split_clause, [], [f8693, f1291, f595, f565])).
fof(f1291, plain, (spl30_192 <=> (e2 = op(op(e0, e2), e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_192])])).
fof(f8693, plain, (~ (e2 = op(e1, e0)) | (~ spl30_54 | spl30_192)), inference(forward_demodulation, [], [f1292, f597])).
fof(f1292, plain, (~ (e2 = op(op(e0, e2), e0)) | spl30_192), inference(avatar_component_clause, [], [f1291])).
fof(f8692, plain, (~ spl30_9 | ~ spl30_60 | spl30_195), inference(avatar_split_clause, [], [f8691, f1305, f620, f404])).
fof(f620, plain, (spl30_60 <=> (e3 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_60])])).
fof(f1305, plain, (spl30_195 <=> (e0 = op(op(e0, e1), e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_195])])).
fof(f8691, plain, (~ (e0 = op(e3, e1)) | (~ spl30_60 | spl30_195)), inference(forward_demodulation, [], [f1307, f622])).
fof(f622, plain, ((e3 = op(e0, e1)) | ~ spl30_60), inference(avatar_component_clause, [], [f620])).
fof(f1307, plain, (~ (e0 = op(op(e0, e1), e1)) | spl30_195), inference(avatar_component_clause, [], [f1305])).
fof(f8681, plain, (spl30_48 | ~ spl30_60 | ~ spl30_125), inference(avatar_split_clause, [], [f8680, f966, f620, f569])).
fof(f966, plain, (spl30_125 <=> (op(e0, e1) = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_125])])).
fof(f8680, plain, ((e3 = op(e1, e0)) | (~ spl30_60 | ~ spl30_125)), inference(forward_demodulation, [], [f968, f622])).
fof(f968, plain, ((op(e0, e1) = op(e1, e0)) | ~ spl30_125), inference(avatar_component_clause, [], [f966])).
fof(f8673, plain, (~ spl30_8 | ~ spl30_51 | spl30_127), inference(avatar_split_clause, [], [f8345, f975, f582, f399])).
fof(f975, plain, (spl30_127 <=> (e3 = op(e3, op(e0, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl30_127])])).
fof(f8345, plain, (~ (e3 = op(e3, e2)) | (~ spl30_51 | spl30_127)), inference(backward_demodulation, [], [f977, f584])).
fof(f977, plain, (~ (e3 = op(e3, op(e0, e3))) | spl30_127), inference(avatar_component_clause, [], [f975])).
fof(f8668, plain, (~ spl30_47 | spl30_59 | ~ spl30_125), inference(avatar_contradiction_clause, [], [f8667])).
fof(f8667, plain, ($false | (~ spl30_47 | spl30_59 | ~ spl30_125)), inference(subsumption_resolution, [], [f8666, f617])).
fof(f617, plain, (~ (e2 = op(e0, e1)) | spl30_59), inference(avatar_component_clause, [], [f616])).
fof(f616, plain, (spl30_59 <=> (e2 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_59])])).
fof(f8666, plain, ((e2 = op(e0, e1)) | (~ spl30_47 | ~ spl30_125)), inference(forward_demodulation, [], [f968, f567])).
fof(f567, plain, ((e2 = op(e1, e0)) | ~ spl30_47), inference(avatar_component_clause, [], [f565])).
fof(f8665, plain, (spl30_14 | ~ spl30_60 | ~ spl30_196), inference(avatar_contradiction_clause, [], [f8664])).
fof(f8664, plain, ($false | (spl30_14 | ~ spl30_60 | ~ spl30_196)), inference(subsumption_resolution, [], [f8663, f426])).
fof(f426, plain, (~ (e1 = op(e3, e0)) | spl30_14), inference(avatar_component_clause, [], [f425])).
fof(f425, plain, (spl30_14 <=> (e1 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_14])])).
fof(f8663, plain, ((e1 = op(e3, e0)) | (~ spl30_60 | ~ spl30_196)), inference(forward_demodulation, [], [f1312, f622])).
fof(f1312, plain, ((e1 = op(op(e0, e1), e0)) | ~ spl30_196), inference(avatar_component_clause, [], [f1310])).
fof(f1310, plain, (spl30_196 <=> (e1 = op(op(e0, e1), e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_196])])).
fof(f8648, plain, (~ spl30_4 | ~ spl30_16), inference(avatar_split_clause, [], [f8642, f433, f382])).
fof(f382, plain, (spl30_4 <=> (e3 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_4])])).
fof(f8642, plain, (~ (e3 = op(e3, e3)) | ~ spl30_16), inference(backward_demodulation, [], [f207, f435])).
fof(f435, plain, ((e3 = op(e3, e0)) | ~ spl30_16), inference(avatar_component_clause, [], [f433])).
fof(f207, plain, ~ (op(e3, e0) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f8638, plain, (~ spl30_4 | ~ spl30_20), inference(avatar_split_clause, [], [f8630, f450, f382])).
fof(f8630, plain, (~ (e3 = op(e3, e3)) | ~ spl30_20), inference(backward_demodulation, [], [f186, f452])).
fof(f186, plain, ~ (op(e2, e3) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f8629, plain, (~ spl30_17 | ~ spl30_21), inference(avatar_split_clause, [], [f8622, f455, f438])).
fof(f438, plain, (spl30_17 <=> (e0 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_17])])).
fof(f455, plain, (spl30_21 <=> (e0 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_21])])).
fof(f8622, plain, (~ (e0 = op(e2, e3)) | ~ spl30_21), inference(backward_demodulation, [], [f204, f457])).
fof(f457, plain, ((e0 = op(e2, e2)) | ~ spl30_21), inference(avatar_component_clause, [], [f455])).
fof(f204, plain, ~ (op(e2, e2) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f8628, plain, (~ spl30_5 | ~ spl30_21), inference(avatar_split_clause, [], [f8621, f455, f387])).
fof(f387, plain, (spl30_5 <=> (e0 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_5])])).
fof(f8621, plain, (~ (e0 = op(e3, e2)) | ~ spl30_21), inference(backward_demodulation, [], [f180, f457])).
fof(f180, plain, ~ (op(e2, e2) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f8613, plain, (~ spl30_30 | ~ spl30_42 | spl30_102), inference(avatar_contradiction_clause, [], [f8612])).
fof(f8612, plain, ($false | (~ spl30_30 | ~ spl30_42 | spl30_102)), inference(subsumption_resolution, [], [f8604, f546])).
fof(f546, plain, ((e1 = op(e1, e1)) | ~ spl30_42), inference(avatar_component_clause, [], [f544])).
fof(f8604, plain, (~ (e1 = op(e1, e1)) | (~ spl30_30 | spl30_102)), inference(backward_demodulation, [], [f855, f495])).
fof(f855, plain, (~ (e1 = op(e1, op(e2, e0))) | spl30_102), inference(avatar_component_clause, [], [f853])).
fof(f853, plain, (spl30_102 <=> (e1 = op(e1, op(e2, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl30_102])])).
fof(f8611, plain, (~ spl30_27 | ~ spl30_30 | spl30_101), inference(avatar_split_clause, [], [f8603, f848, f493, f480])).
fof(f480, plain, (spl30_27 <=> (e2 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_27])])).
fof(f848, plain, (spl30_101 <=> (e2 = op(e2, op(e2, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl30_101])])).
fof(f8603, plain, (~ (e2 = op(e2, e1)) | (~ spl30_30 | spl30_101)), inference(backward_demodulation, [], [f850, f495])).
fof(f850, plain, (~ (e2 = op(e2, op(e2, e0))) | spl30_101), inference(avatar_component_clause, [], [f848])).
fof(f8610, plain, (~ spl30_26 | ~ spl30_30), inference(avatar_split_clause, [], [f8597, f493, f476])).
fof(f476, plain, (spl30_26 <=> (e1 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_26])])).
fof(f8597, plain, (~ (e1 = op(e2, e1)) | ~ spl30_30), inference(backward_demodulation, [], [f199, f495])).
fof(f199, plain, ~ (op(e2, e0) = op(e2, e1)), inference(cnf_transformation, [], [f3])).
fof(f8609, plain, (~ spl30_14 | ~ spl30_30), inference(avatar_split_clause, [], [f8596, f493, f425])).
fof(f8596, plain, (~ (e1 = op(e3, e0)) | ~ spl30_30), inference(backward_demodulation, [], [f168, f495])).
fof(f168, plain, ~ (op(e2, e0) = op(e3, e0)), inference(cnf_transformation, [], [f3])).
fof(f8590, plain, (spl30_9 | ~ spl30_33 | ~ spl30_76), inference(avatar_split_clause, [], [f8577, f727, f506, f404])).
fof(f8577, plain, ((e0 = op(e3, e1)) | (~ spl30_33 | ~ spl30_76)), inference(backward_demodulation, [], [f729, f508])).
fof(f508, plain, ((e0 = op(e1, e3)) | ~ spl30_33), inference(avatar_component_clause, [], [f506])).
fof(f8589, plain, (~ spl30_17 | ~ spl30_33), inference(avatar_split_clause, [], [f8572, f506, f438])).
fof(f8572, plain, (~ (e0 = op(e2, e3)) | ~ spl30_33), inference(backward_demodulation, [], [f184, f508])).
fof(f184, plain, ~ (op(e1, e3) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f8570, plain, (~ spl30_34 | ~ spl30_40 | spl30_113), inference(avatar_split_clause, [], [f8563, f907, f535, f510])).
fof(f907, plain, (spl30_113 <=> (e1 = op(e1, op(e1, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl30_113])])).
fof(f8563, plain, (~ (e1 = op(e1, e3)) | (~ spl30_40 | spl30_113)), inference(backward_demodulation, [], [f909, f537])).
fof(f537, plain, ((e3 = op(e1, e2)) | ~ spl30_40), inference(avatar_component_clause, [], [f535])).
fof(f909, plain, (~ (e1 = op(e1, op(e1, e2))) | spl30_113), inference(avatar_component_clause, [], [f907])).
fof(f8569, plain, (~ spl30_4 | ~ spl30_40 | spl30_111), inference(avatar_split_clause, [], [f8561, f897, f535, f382])).
fof(f897, plain, (spl30_111 <=> (e3 = op(e3, op(e1, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl30_111])])).
fof(f8561, plain, (~ (e3 = op(e3, e3)) | (~ spl30_40 | spl30_111)), inference(backward_demodulation, [], [f899, f537])).
fof(f899, plain, (~ (e3 = op(e3, op(e1, e2))) | spl30_111), inference(avatar_component_clause, [], [f897])).
fof(f8568, plain, (~ spl30_24 | ~ spl30_40), inference(avatar_split_clause, [], [f8557, f535, f467])).
fof(f467, plain, (spl30_24 <=> (e3 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_24])])).
fof(f8557, plain, (~ (e3 = op(e2, e2)) | ~ spl30_40), inference(backward_demodulation, [], [f178, f537])).
fof(f178, plain, ~ (op(e1, e2) = op(e2, e2)), inference(cnf_transformation, [], [f3])).
fof(f8554, plain, (~ spl30_34 | ~ spl30_42), inference(avatar_split_clause, [], [f8545, f544, f510])).
fof(f8545, plain, (~ (e1 = op(e1, e3)) | ~ spl30_42), inference(backward_demodulation, [], [f197, f546])).
fof(f8553, plain, (~ spl30_26 | ~ spl30_42), inference(avatar_split_clause, [], [f8542, f544, f476])).
fof(f8542, plain, (~ (e1 = op(e2, e1)) | ~ spl30_42), inference(backward_demodulation, [], [f172, f546])).
fof(f172, plain, ~ (op(e1, e1) = op(e2, e1)), inference(cnf_transformation, [], [f3])).
fof(f8538, plain, (~ spl30_39 | ~ spl30_47), inference(avatar_split_clause, [], [f8527, f565, f531])).
fof(f8527, plain, (~ (e2 = op(e1, e2)) | ~ spl30_47), inference(backward_demodulation, [], [f194, f567])).
fof(f8537, plain, (~ spl30_31 | ~ spl30_47), inference(avatar_split_clause, [], [f8524, f565, f497])).
fof(f497, plain, (spl30_31 <=> (e2 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_31])])).
fof(f8524, plain, (~ (e2 = op(e2, e0)) | ~ spl30_47), inference(backward_demodulation, [], [f166, f567])).
fof(f166, plain, ~ (op(e1, e0) = op(e2, e0)), inference(cnf_transformation, [], [f3])).
fof(f8511, plain, (spl30_30 | ~ spl30_54 | ~ spl30_104), inference(avatar_split_clause, [], [f8510, f863, f595, f493])).
fof(f863, plain, (spl30_104 <=> (op(e0, e2) = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_104])])).
fof(f8510, plain, ((e1 = op(e2, e0)) | (~ spl30_54 | ~ spl30_104)), inference(forward_demodulation, [], [f865, f597])).
fof(f865, plain, ((op(e0, e2) = op(e2, e0)) | ~ spl30_104), inference(avatar_component_clause, [], [f863])).
fof(f8504, plain, (~ spl30_48 | ~ spl30_60 | spl30_125), inference(avatar_split_clause, [], [f8503, f966, f620, f569])).
fof(f8503, plain, (~ (e3 = op(e1, e0)) | (~ spl30_60 | spl30_125)), inference(forward_demodulation, [], [f967, f622])).
fof(f967, plain, (~ (op(e0, e1) = op(e1, e0)) | spl30_125), inference(avatar_component_clause, [], [f966])).
fof(f8500, plain, (spl30_42 | ~ spl30_54 | ~ spl30_134), inference(avatar_split_clause, [], [f8499, f1010, f595, f544])).
fof(f8499, plain, ((e1 = op(e1, e1)) | (~ spl30_54 | ~ spl30_134)), inference(forward_demodulation, [], [f1011, f597])).
fof(f1011, plain, ((e1 = op(e1, op(e0, e2))) | ~ spl30_134), inference(avatar_component_clause, [], [f1010])).
fof(f8498, plain, (~ spl30_4 | ~ spl30_60 | spl30_137), inference(avatar_split_clause, [], [f8497, f1025, f620, f382])).
fof(f1025, plain, (spl30_137 <=> (e3 = op(e3, op(e0, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl30_137])])).
fof(f8497, plain, (~ (e3 = op(e3, e3)) | (~ spl30_60 | spl30_137)), inference(forward_demodulation, [], [f1027, f622])).
fof(f1027, plain, (~ (e3 = op(e3, op(e0, e1))) | spl30_137), inference(avatar_component_clause, [], [f1025])).
fof(f8491, plain, (spl30_47 | ~ spl30_54 | ~ spl30_192), inference(avatar_split_clause, [], [f8490, f1291, f595, f565])).
fof(f8490, plain, ((e2 = op(e1, e0)) | (~ spl30_54 | ~ spl30_192)), inference(forward_demodulation, [], [f1293, f597])).
fof(f1293, plain, ((e2 = op(op(e0, e2), e0)) | ~ spl30_192), inference(avatar_component_clause, [], [f1291])).
fof(f8488, plain, (spl30_9 | ~ spl30_60 | ~ spl30_195), inference(avatar_split_clause, [], [f8487, f1305, f620, f404])).
fof(f8487, plain, ((e0 = op(e3, e1)) | (~ spl30_60 | ~ spl30_195)), inference(forward_demodulation, [], [f1306, f622])).
fof(f1306, plain, ((e0 = op(op(e0, e1), e1)) | ~ spl30_195), inference(avatar_component_clause, [], [f1305])).
fof(f8479, plain, (~ spl30_14 | ~ spl30_223 | spl30_225), inference(avatar_split_clause, [], [f8477, f1457, f1448, f425])).
fof(f1448, plain, (spl30_223 <=> (e0 = op(e3, op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl30_223])])).
fof(f1457, plain, (spl30_225 <=> (e1 = op(e3, op(e3, op(e3, e3))))), introduced(avatar_definition, [new_symbols(naming, [spl30_225])])).
fof(f8477, plain, (~ (e1 = op(e3, e0)) | (~ spl30_223 | spl30_225)), inference(backward_demodulation, [], [f1459, f1449])).
fof(f1449, plain, ((e0 = op(e3, op(e3, e3))) | ~ spl30_223), inference(avatar_component_clause, [], [f1448])).
fof(f1459, plain, (~ (e1 = op(e3, op(e3, op(e3, e3)))) | spl30_225), inference(avatar_component_clause, [], [f1457])).
fof(f8473, plain, (~ spl30_4 | spl30_49 | ~ spl30_203), inference(avatar_contradiction_clause, [], [f8472])).
fof(f8472, plain, ($false | (~ spl30_4 | spl30_49 | ~ spl30_203)), inference(subsumption_resolution, [], [f8467, f575])).
fof(f575, plain, (~ (e0 = op(e0, e3)) | spl30_49), inference(avatar_component_clause, [], [f574])).
fof(f574, plain, (spl30_49 <=> (e0 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_49])])).
fof(f8467, plain, ((e0 = op(e0, e3)) | (~ spl30_4 | ~ spl30_203)), inference(backward_demodulation, [], [f1347, f384])).
fof(f384, plain, ((e3 = op(e3, e3)) | ~ spl30_4), inference(avatar_component_clause, [], [f382])).
fof(f1347, plain, ((e0 = op(e0, op(e3, e3))) | ~ spl30_203), inference(avatar_component_clause, [], [f1346])).
fof(f1346, plain, (spl30_203 <=> (e0 = op(e0, op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl30_203])])).
fof(f8446, plain, (spl30_5 | ~ spl30_17 | ~ spl30_70), inference(avatar_split_clause, [], [f8437, f698, f438, f387])).
fof(f8437, plain, ((e0 = op(e3, e2)) | (~ spl30_17 | ~ spl30_70)), inference(backward_demodulation, [], [f700, f440])).
fof(f440, plain, ((e0 = op(e2, e3)) | ~ spl30_17), inference(avatar_component_clause, [], [f438])).
fof(f8445, plain, (~ spl30_1 | ~ spl30_17), inference(avatar_split_clause, [], [f8433, f438, f370])).
fof(f370, plain, (spl30_1 <=> (e0 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_1])])).
fof(f8433, plain, (~ (e0 = op(e3, e3)) | ~ spl30_17), inference(backward_demodulation, [], [f186, f440])).
fof(f8430, plain, (~ spl30_20 | ~ spl30_24), inference(avatar_split_clause, [], [f8426, f467, f450])).
fof(f8426, plain, (~ (e3 = op(e2, e3)) | ~ spl30_24), inference(backward_demodulation, [], [f204, f469])).
fof(f469, plain, ((e3 = op(e2, e2)) | ~ spl30_24), inference(avatar_component_clause, [], [f467])).
fof(f8421, plain, (~ spl30_26 | spl30_57 | ~ spl30_97), inference(avatar_contradiction_clause, [], [f8420])).
fof(f8420, plain, ($false | (~ spl30_26 | spl30_57 | ~ spl30_97)), inference(subsumption_resolution, [], [f8413, f609])).
fof(f609, plain, (~ (e0 = op(e0, e1)) | spl30_57), inference(avatar_component_clause, [], [f608])).
fof(f608, plain, (spl30_57 <=> (e0 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_57])])).
fof(f8413, plain, ((e0 = op(e0, e1)) | (~ spl30_26 | ~ spl30_97)), inference(backward_demodulation, [], [f830, f478])).
fof(f478, plain, ((e1 = op(e2, e1)) | ~ spl30_26), inference(avatar_component_clause, [], [f476])).
fof(f830, plain, ((e0 = op(e0, op(e2, e1))) | ~ spl30_97), inference(avatar_component_clause, [], [f829])).
fof(f829, plain, (spl30_97 <=> (e0 = op(e0, op(e2, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl30_97])])).
fof(f8418, plain, (~ spl30_12 | ~ spl30_26 | spl30_94), inference(avatar_split_clause, [], [f8411, f814, f476, f416])).
fof(f416, plain, (spl30_12 <=> (e3 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_12])])).
fof(f814, plain, (spl30_94 <=> (e3 = op(e3, op(e2, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl30_94])])).
fof(f8411, plain, (~ (e3 = op(e3, e1)) | (~ spl30_26 | spl30_94)), inference(backward_demodulation, [], [f816, f478])).
fof(f816, plain, (~ (e3 = op(e3, op(e2, e1))) | spl30_94), inference(avatar_component_clause, [], [f814])).
fof(f8417, plain, (~ spl30_22 | ~ spl30_26), inference(avatar_split_clause, [], [f8407, f476, f459])).
fof(f459, plain, (spl30_22 <=> (e1 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_22])])).
fof(f8407, plain, (~ (e1 = op(e2, e2)) | ~ spl30_26), inference(backward_demodulation, [], [f202, f478])).
fof(f202, plain, ~ (op(e2, e1) = op(e2, e2)), inference(cnf_transformation, [], [f3])).
fof(f8387, plain, (~ spl30_35 | ~ spl30_39), inference(avatar_split_clause, [], [f8376, f531, f514])).
fof(f514, plain, (spl30_35 <=> (e2 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_35])])).
fof(f8376, plain, (~ (e2 = op(e1, e3)) | ~ spl30_39), inference(backward_demodulation, [], [f198, f533])).
fof(f198, plain, ~ (op(e1, e2) = op(e1, e3)), inference(cnf_transformation, [], [f3])).
fof(f8372, plain, (~ spl30_37 | ~ spl30_41), inference(avatar_split_clause, [], [f8359, f540, f523])).
fof(f523, plain, (spl30_37 <=> (e0 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_37])])).
fof(f8359, plain, (~ (e0 = op(e1, e2)) | ~ spl30_41), inference(backward_demodulation, [], [f196, f542])).
fof(f196, plain, ~ (op(e1, e1) = op(e1, e2)), inference(cnf_transformation, [], [f3])).
fof(f8371, plain, (~ spl30_25 | ~ spl30_41), inference(avatar_split_clause, [], [f8357, f540, f472])).
fof(f472, plain, (spl30_25 <=> (e0 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_25])])).
fof(f8357, plain, (~ (e0 = op(e2, e1)) | ~ spl30_41), inference(backward_demodulation, [], [f172, f542])).
fof(f8355, plain, (~ spl30_31 | ~ spl30_51 | ~ spl30_188), inference(avatar_contradiction_clause, [], [f8354])).
fof(f8354, plain, ($false | (~ spl30_31 | ~ spl30_51 | ~ spl30_188)), inference(subsumption_resolution, [], [f8353, f216])).
fof(f216, plain, ~ (e2 = e3), inference(cnf_transformation, [], [f4])).
fof(f8353, plain, ((e2 = e3) | (~ spl30_31 | ~ spl30_51 | ~ spl30_188)), inference(forward_demodulation, [], [f8349, f499])).
fof(f499, plain, ((e2 = op(e2, e0)) | ~ spl30_31), inference(avatar_component_clause, [], [f497])).
fof(f8349, plain, ((e3 = op(e2, e0)) | (~ spl30_51 | ~ spl30_188)), inference(backward_demodulation, [], [f1274, f584])).
fof(f8352, plain, (~ spl30_17 | ~ spl30_51 | spl30_187), inference(avatar_split_clause, [], [f8348, f1267, f582, f438])).
fof(f1267, plain, (spl30_187 <=> (e0 = op(op(e0, e3), e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_187])])).
fof(f8348, plain, (~ (e0 = op(e2, e3)) | (~ spl30_51 | spl30_187)), inference(backward_demodulation, [], [f1269, f584])).
fof(f1269, plain, (~ (e0 = op(op(e0, e3), e3)) | spl30_187), inference(avatar_component_clause, [], [f1267])).
fof(f8351, plain, (~ spl30_35 | ~ spl30_51), inference(avatar_split_clause, [], [f8341, f582, f514])).
fof(f8341, plain, (~ (e2 = op(e1, e3)) | ~ spl30_51), inference(backward_demodulation, [], [f181, f584])).
fof(f181, plain, ~ (op(e0, e3) = op(e1, e3)), inference(cnf_transformation, [], [f3])).
fof(f8340, plain, (~ spl30_37 | ~ spl30_54 | spl30_191), inference(avatar_split_clause, [], [f8335, f1286, f595, f523])).
fof(f1286, plain, (spl30_191 <=> (e0 = op(op(e0, e2), e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_191])])).
fof(f8335, plain, (~ (e0 = op(e1, e2)) | (~ spl30_54 | spl30_191)), inference(backward_demodulation, [], [f1288, f597])).
fof(f1288, plain, (~ (e0 = op(op(e0, e2), e2)) | spl30_191), inference(avatar_component_clause, [], [f1286])).
fof(f8338, plain, (~ spl30_50 | ~ spl30_54), inference(avatar_split_clause, [], [f8330, f595, f578])).
fof(f578, plain, (spl30_50 <=> (e1 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_50])])).
fof(f8330, plain, (~ (e1 = op(e0, e3)) | ~ spl30_54), inference(backward_demodulation, [], [f192, f597])).
fof(f192, plain, ~ (op(e0, e2) = op(e0, e3)), inference(cnf_transformation, [], [f3])).
fof(f8337, plain, (~ spl30_22 | ~ spl30_54), inference(avatar_split_clause, [], [f8328, f595, f459])).
fof(f8328, plain, (~ (e1 = op(e2, e2)) | ~ spl30_54), inference(backward_demodulation, [], [f176, f597])).
fof(f176, plain, ~ (op(e0, e2) = op(e2, e2)), inference(cnf_transformation, [], [f3])).
fof(f8326, plain, (~ spl30_50 | ~ spl30_60 | spl30_197), inference(avatar_split_clause, [], [f8321, f1315, f620, f578])).
fof(f1315, plain, (spl30_197 <=> (e1 = op(e0, op(e0, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl30_197])])).
fof(f8321, plain, (~ (e1 = op(e0, e3)) | (~ spl30_60 | spl30_197)), inference(backward_demodulation, [], [f1317, f622])).
fof(f1317, plain, (~ (e1 = op(e0, op(e0, e1))) | spl30_197), inference(avatar_component_clause, [], [f1315])).
fof(f8324, plain, (spl30_4 | ~ spl30_60 | ~ spl30_137), inference(avatar_split_clause, [], [f8315, f1025, f620, f382])).
fof(f8315, plain, ((e3 = op(e3, e3)) | (~ spl30_60 | ~ spl30_137)), inference(backward_demodulation, [], [f1026, f622])).
fof(f1026, plain, ((e3 = op(e3, op(e0, e1))) | ~ spl30_137), inference(avatar_component_clause, [], [f1025])).
fof(f8323, plain, (~ spl30_56 | ~ spl30_60), inference(avatar_split_clause, [], [f8313, f620, f603])).
fof(f603, plain, (spl30_56 <=> (e3 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_56])])).
fof(f8313, plain, (~ (e3 = op(e0, e2)) | ~ spl30_60), inference(backward_demodulation, [], [f190, f622])).
fof(f190, plain, ~ (op(e0, e1) = op(e0, e2)), inference(cnf_transformation, [], [f3])).
fof(f8322, plain, (~ spl30_12 | ~ spl30_60), inference(avatar_split_clause, [], [f8312, f620, f416])).
fof(f8312, plain, (~ (e3 = op(e3, e1)) | ~ spl30_60), inference(backward_demodulation, [], [f171, f622])).
fof(f171, plain, ~ (op(e0, e1) = op(e3, e1)), inference(cnf_transformation, [], [f3])).
fof(f8299, plain, (~ spl30_12 | ~ spl30_14 | spl30_78), inference(avatar_split_clause, [], [f8298, f736, f425, f416])).
fof(f736, plain, (spl30_78 <=> (e3 = op(e3, op(e3, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl30_78])])).
fof(f8298, plain, (~ (e3 = op(e3, e1)) | (~ spl30_14 | spl30_78)), inference(forward_demodulation, [], [f738, f427])).
fof(f427, plain, ((e1 = op(e3, e0)) | ~ spl30_14), inference(avatar_component_clause, [], [f425])).
fof(f738, plain, (~ (e3 = op(e3, op(e3, e0))) | spl30_78), inference(avatar_component_clause, [], [f736])).
fof(f8296, plain, (~ spl30_42 | ~ spl30_14 | spl30_80), inference(avatar_split_clause, [], [f8295, f746, f425, f544])).
fof(f746, plain, (spl30_80 <=> (e1 = op(e1, op(e3, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl30_80])])).
fof(f8295, plain, (~ (e1 = op(e1, e1)) | (~ spl30_14 | spl30_80)), inference(forward_demodulation, [], [f748, f427])).
fof(f748, plain, (~ (e1 = op(e1, op(e3, e0))) | spl30_80), inference(avatar_component_clause, [], [f746])).
fof(f8289, plain, (spl30_60 | ~ spl30_48 | ~ spl30_125), inference(avatar_split_clause, [], [f8288, f966, f569, f620])).
fof(f8288, plain, ((e3 = op(e0, e1)) | (~ spl30_48 | ~ spl30_125)), inference(forward_demodulation, [], [f968, f571])).
fof(f8283, plain, (~ spl30_35 | spl30_213 | ~ spl30_216), inference(avatar_split_clause, [], [f8281, f1412, f1397, f514])).
fof(f1397, plain, (spl30_213 <=> (e2 = op(e1, op(e1, op(e1, e1))))), introduced(avatar_definition, [new_symbols(naming, [spl30_213])])).
fof(f1412, plain, (spl30_216 <=> (e3 = op(e1, op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl30_216])])).
fof(f8281, plain, (~ (e2 = op(e1, e3)) | (spl30_213 | ~ spl30_216)), inference(backward_demodulation, [], [f1399, f1413])).
fof(f1413, plain, ((e3 = op(e1, op(e1, e1))) | ~ spl30_216), inference(avatar_component_clause, [], [f1412])).
fof(f1399, plain, (~ (e2 = op(e1, op(e1, op(e1, e1)))) | spl30_213), inference(avatar_component_clause, [], [f1397])).
fof(f8258, plain, (~ spl30_9 | ~ spl30_14 | spl30_157), inference(avatar_split_clause, [], [f8255, f1122, f425, f404])).
fof(f1122, plain, (spl30_157 <=> (e0 = op(e3, op(e3, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl30_157])])).
fof(f8255, plain, (~ (e0 = op(e3, e1)) | (~ spl30_14 | spl30_157)), inference(backward_demodulation, [], [f1124, f427])).
fof(f1124, plain, (~ (e0 = op(e3, op(e3, e0))) | spl30_157), inference(avatar_component_clause, [], [f1122])).
fof(f8256, plain, (~ spl30_6 | ~ spl30_14), inference(avatar_split_clause, [], [f8246, f425, f391])).
fof(f391, plain, (spl30_6 <=> (e1 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_6])])).
fof(f8246, plain, (~ (e1 = op(e3, e2)) | ~ spl30_14), inference(backward_demodulation, [], [f206, f427])).
fof(f8233, plain, (~ spl30_6 | ~ spl30_22), inference(avatar_split_clause, [], [f8226, f459, f391])).
fof(f8226, plain, (~ (e1 = op(e3, e2)) | ~ spl30_22), inference(backward_demodulation, [], [f180, f461])).
fof(f461, plain, ((e1 = op(e2, e2)) | ~ spl30_22), inference(avatar_component_clause, [], [f459])).
fof(f8220, plain, (~ spl30_19 | ~ spl30_31), inference(avatar_split_clause, [], [f8213, f497, f446])).
fof(f446, plain, (spl30_19 <=> (e2 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_19])])).
fof(f8213, plain, (~ (e2 = op(e2, e3)) | ~ spl30_31), inference(backward_demodulation, [], [f201, f499])).
fof(f201, plain, ~ (op(e2, e0) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f8219, plain, (~ spl30_15 | ~ spl30_31), inference(avatar_split_clause, [], [f8211, f497, f429])).
fof(f8211, plain, (~ (e2 = op(e3, e0)) | ~ spl30_31), inference(backward_demodulation, [], [f168, f499])).
fof(f8206, plain, (spl30_11 | ~ spl30_35 | ~ spl30_76), inference(avatar_contradiction_clause, [], [f8205])).
fof(f8205, plain, ($false | (spl30_11 | ~ spl30_35 | ~ spl30_76)), inference(subsumption_resolution, [], [f8194, f413])).
fof(f413, plain, (~ (e2 = op(e3, e1)) | spl30_11), inference(avatar_component_clause, [], [f412])).
fof(f412, plain, (spl30_11 <=> (e2 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_11])])).
fof(f8194, plain, ((e2 = op(e3, e1)) | (~ spl30_35 | ~ spl30_76)), inference(backward_demodulation, [], [f729, f516])).
fof(f516, plain, ((e2 = op(e1, e3)) | ~ spl30_35), inference(avatar_component_clause, [], [f514])).
fof(f8204, plain, (~ spl30_19 | ~ spl30_35), inference(avatar_split_clause, [], [f8191, f514, f446])).
fof(f8191, plain, (~ (e2 = op(e2, e3)) | ~ spl30_35), inference(backward_demodulation, [], [f184, f516])).
fof(f8189, plain, (spl30_25 | ~ spl30_37 | ~ spl30_98), inference(avatar_split_clause, [], [f8176, f834, f523, f472])).
fof(f834, plain, (spl30_98 <=> (op(e1, e2) = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_98])])).
fof(f8176, plain, ((e0 = op(e2, e1)) | (~ spl30_37 | ~ spl30_98)), inference(backward_demodulation, [], [f836, f525])).
fof(f525, plain, ((e0 = op(e1, e2)) | ~ spl30_37), inference(avatar_component_clause, [], [f523])).
fof(f836, plain, ((op(e1, e2) = op(e2, e1)) | ~ spl30_98), inference(avatar_component_clause, [], [f834])).
fof(f8188, plain, (~ spl30_33 | ~ spl30_37), inference(avatar_split_clause, [], [f8172, f523, f506])).
fof(f8172, plain, (~ (e0 = op(e1, e3)) | ~ spl30_37), inference(backward_demodulation, [], [f198, f525])).
fof(f8187, plain, (~ spl30_21 | ~ spl30_37), inference(avatar_split_clause, [], [f8170, f523, f455])).
fof(f8170, plain, (~ (e0 = op(e2, e2)) | ~ spl30_37), inference(backward_demodulation, [], [f178, f525])).
fof(f8153, plain, (~ spl30_44 | ~ spl30_48), inference(avatar_split_clause, [], [f8143, f569, f552])).
fof(f552, plain, (spl30_44 <=> (e3 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_44])])).
fof(f8143, plain, (~ (e3 = op(e1, e1)) | ~ spl30_48), inference(backward_demodulation, [], [f193, f571])).
fof(f193, plain, ~ (op(e1, e0) = op(e1, e1)), inference(cnf_transformation, [], [f3])).
fof(f8152, plain, (~ spl30_32 | ~ spl30_48), inference(avatar_split_clause, [], [f8141, f569, f501])).
fof(f501, plain, (spl30_32 <=> (e3 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_32])])).
fof(f8141, plain, (~ (e3 = op(e2, e0)) | ~ spl30_48), inference(backward_demodulation, [], [f166, f571])).
fof(f8138, plain, (spl30_14 | ~ spl30_50 | ~ spl30_82), inference(avatar_split_clause, [], [f8137, f756, f578, f425])).
fof(f8137, plain, ((e1 = op(e3, e0)) | (~ spl30_50 | ~ spl30_82)), inference(forward_demodulation, [], [f758, f580])).
fof(f580, plain, ((e1 = op(e0, e3)) | ~ spl30_50), inference(avatar_component_clause, [], [f578])).
fof(f8125, plain, (~ spl30_32 | ~ spl30_56 | spl30_104), inference(avatar_split_clause, [], [f8124, f863, f603, f501])).
fof(f8124, plain, (~ (e3 = op(e2, e0)) | (~ spl30_56 | spl30_104)), inference(forward_demodulation, [], [f864, f605])).
fof(f605, plain, ((e3 = op(e0, e2)) | ~ spl30_56), inference(avatar_component_clause, [], [f603])).
fof(f864, plain, (~ (op(e0, e2) = op(e2, e0)) | spl30_104), inference(avatar_component_clause, [], [f863])).
fof(f8112, plain, (~ spl30_46 | ~ spl30_61 | spl30_144), inference(avatar_split_clause, [], [f8111, f1060, f625, f561])).
fof(f561, plain, (spl30_46 <=> (e1 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_46])])).
fof(f1060, plain, (spl30_144 <=> (e1 = op(e1, op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl30_144])])).
fof(f8111, plain, (~ (e1 = op(e1, e0)) | (~ spl30_61 | spl30_144)), inference(forward_demodulation, [], [f1062, f627])).
fof(f1062, plain, (~ (e1 = op(e1, op(e0, e0))) | spl30_144), inference(avatar_component_clause, [], [f1060])).
fof(f8105, plain, (~ spl30_33 | ~ spl30_50 | spl30_187), inference(avatar_split_clause, [], [f8104, f1267, f578, f506])).
fof(f8104, plain, (~ (e0 = op(e1, e3)) | (~ spl30_50 | spl30_187)), inference(forward_demodulation, [], [f1269, f580])).
fof(f8103, plain, (spl30_48 | ~ spl30_50 | ~ spl30_188), inference(avatar_split_clause, [], [f8102, f1272, f578, f569])).
fof(f8102, plain, ((e3 = op(e1, e0)) | (~ spl30_50 | ~ spl30_188)), inference(forward_demodulation, [], [f1274, f580])).
fof(f8101, plain, (~ spl30_15 | ~ spl30_56 | spl30_192), inference(avatar_split_clause, [], [f8100, f1291, f603, f429])).
fof(f8100, plain, (~ (e2 = op(e3, e0)) | (~ spl30_56 | spl30_192)), inference(forward_demodulation, [], [f1292, f605])).
fof(f8099, plain, (spl30_25 | ~ spl30_59 | ~ spl30_195), inference(avatar_split_clause, [], [f8098, f1305, f616, f472])).
fof(f8098, plain, ((e0 = op(e2, e1)) | (~ spl30_59 | ~ spl30_195)), inference(forward_demodulation, [], [f1306, f618])).
fof(f618, plain, ((e2 = op(e0, e1)) | ~ spl30_59), inference(avatar_component_clause, [], [f616])).
fof(f8096, plain, (~ spl30_32 | ~ spl30_217 | spl30_218), inference(avatar_split_clause, [], [f8092, f1422, f1418, f501])).
fof(f1418, plain, (spl30_217 <=> (e0 = op(e2, op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl30_217])])).
fof(f1422, plain, (spl30_218 <=> (e3 = op(e2, op(e2, op(e2, e2))))), introduced(avatar_definition, [new_symbols(naming, [spl30_218])])).
fof(f8092, plain, (~ (e3 = op(e2, e0)) | (~ spl30_217 | spl30_218)), inference(backward_demodulation, [], [f1424, f1419])).
fof(f1419, plain, ((e0 = op(e2, op(e2, e2))) | ~ spl30_217), inference(avatar_component_clause, [], [f1418])).
fof(f1424, plain, (~ (e3 = op(e2, op(e2, op(e2, e2)))) | spl30_218), inference(avatar_component_clause, [], [f1422])).
fof(f8090, plain, (~ spl30_9 | ~ spl30_226 | spl30_227), inference(avatar_split_clause, [], [f8087, f1467, f1462, f404])).
fof(f1462, plain, (spl30_226 <=> (e1 = op(e3, op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl30_226])])).
fof(f1467, plain, (spl30_227 <=> (e0 = op(e3, op(e3, op(e3, e3))))), introduced(avatar_definition, [new_symbols(naming, [spl30_227])])).
fof(f8087, plain, (~ (e0 = op(e3, e1)) | (~ spl30_226 | spl30_227)), inference(backward_demodulation, [], [f1469, f1463])).
fof(f1463, plain, ((e1 = op(e3, op(e3, e3))) | ~ spl30_226), inference(avatar_component_clause, [], [f1462])).
fof(f1469, plain, (~ (e0 = op(e3, op(e3, op(e3, e3)))) | spl30_227), inference(avatar_component_clause, [], [f1467])).
fof(f8076, plain, (~ spl30_8 | ~ spl30_56), inference(avatar_split_clause, [], [f8075, f603, f399])).
fof(f8075, plain, (~ (e3 = op(e3, e2)) | ~ spl30_56), inference(forward_demodulation, [], [f177, f605])).
fof(f177, plain, ~ (op(e0, e2) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f8033, plain, (~ spl30_23 | ~ spl30_39), inference(avatar_split_clause, [], [f8024, f531, f463])).
fof(f463, plain, (spl30_23 <=> (e2 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_23])])).
fof(f8024, plain, (~ (e2 = op(e2, e2)) | ~ spl30_39), inference(backward_demodulation, [], [f178, f533])).
fof(f8020, plain, (spl30_33 | ~ spl30_9 | ~ spl30_76), inference(avatar_split_clause, [], [f8019, f727, f404, f506])).
fof(f8019, plain, ((e0 = op(e1, e3)) | (~ spl30_9 | ~ spl30_76)), inference(forward_demodulation, [], [f729, f406])).
fof(f406, plain, ((e0 = op(e3, e1)) | ~ spl30_9), inference(avatar_component_clause, [], [f404])).
fof(f8018, plain, (~ spl30_4 | ~ spl30_32 | spl30_100), inference(avatar_contradiction_clause, [], [f8017])).
fof(f8017, plain, ($false | (~ spl30_4 | ~ spl30_32 | spl30_100)), inference(subsumption_resolution, [], [f8016, f384])).
fof(f8016, plain, (~ (e3 = op(e3, e3)) | (~ spl30_32 | spl30_100)), inference(forward_demodulation, [], [f845, f503])).
fof(f503, plain, ((e3 = op(e2, e0)) | ~ spl30_32), inference(avatar_component_clause, [], [f501])).
fof(f845, plain, (~ (e3 = op(e3, op(e2, e0))) | spl30_100), inference(avatar_component_clause, [], [f843])).
fof(f843, plain, (spl30_100 <=> (e3 = op(e3, op(e2, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl30_100])])).
fof(f8015, plain, (~ spl30_19 | ~ spl30_32 | spl30_101), inference(avatar_split_clause, [], [f8014, f848, f501, f446])).
fof(f8014, plain, (~ (e2 = op(e2, e3)) | (~ spl30_32 | spl30_101)), inference(forward_demodulation, [], [f850, f503])).
fof(f8007, plain, (~ spl30_23 | ~ spl30_59 | spl30_138), inference(avatar_split_clause, [], [f8006, f1030, f616, f463])).
fof(f1030, plain, (spl30_138 <=> (e2 = op(e2, op(e0, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl30_138])])).
fof(f8006, plain, (~ (e2 = op(e2, e2)) | (~ spl30_59 | spl30_138)), inference(forward_demodulation, [], [f1032, f618])).
fof(f1032, plain, (~ (e2 = op(e2, op(e0, e1))) | spl30_138), inference(avatar_component_clause, [], [f1030])).
fof(f8000, plain, (~ spl30_17 | ~ spl30_32 | spl30_171), inference(avatar_split_clause, [], [f7999, f1190, f501, f438])).
fof(f1190, plain, (spl30_171 <=> (e0 = op(e2, op(e2, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl30_171])])).
fof(f7999, plain, (~ (e0 = op(e2, e3)) | (~ spl30_32 | spl30_171)), inference(forward_demodulation, [], [f1192, f503])).
fof(f1192, plain, (~ (e0 = op(e2, op(e2, e0))) | spl30_171), inference(avatar_component_clause, [], [f1190])).
fof(f7998, plain, (spl30_33 | ~ spl30_50 | ~ spl30_187), inference(avatar_split_clause, [], [f7997, f1267, f578, f506])).
fof(f7997, plain, ((e0 = op(e1, e3)) | (~ spl30_50 | ~ spl30_187)), inference(forward_demodulation, [], [f1268, f580])).
fof(f1268, plain, ((e0 = op(op(e0, e3), e3)) | ~ spl30_187), inference(avatar_component_clause, [], [f1267])).
fof(f7996, plain, (spl30_19 | ~ spl30_4 | ~ spl30_201), inference(avatar_split_clause, [], [f7995, f1336, f382, f446])).
fof(f1336, plain, (spl30_201 <=> (e2 = op(e2, op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl30_201])])).
fof(f7995, plain, ((e2 = op(e2, e3)) | (~ spl30_4 | ~ spl30_201)), inference(forward_demodulation, [], [f1337, f384])).
fof(f1337, plain, ((e2 = op(e2, op(e3, e3))) | ~ spl30_201), inference(avatar_component_clause, [], [f1336])).
fof(f7987, plain, (~ spl30_40 | ~ spl30_6 | spl30_147), inference(avatar_split_clause, [], [f7650, f1074, f391, f535])).
fof(f1074, plain, (spl30_147 <=> (e3 = op(op(e3, e2), e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_147])])).
fof(f7650, plain, (~ (e3 = op(e1, e2)) | (~ spl30_6 | spl30_147)), inference(backward_demodulation, [], [f1076, f393])).
fof(f393, plain, ((e1 = op(e3, e2)) | ~ spl30_6), inference(avatar_component_clause, [], [f391])).
fof(f1076, plain, (~ (e3 = op(op(e3, e2), e2)) | spl30_147), inference(avatar_component_clause, [], [f1074])).
fof(f7986, plain, (~ spl30_38 | ~ spl30_26 | spl30_98), inference(avatar_split_clause, [], [f7985, f834, f476, f527])).
fof(f527, plain, (spl30_38 <=> (e1 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_38])])).
fof(f7985, plain, (~ (e1 = op(e1, e2)) | (~ spl30_26 | spl30_98)), inference(forward_demodulation, [], [f835, f478])).
fof(f835, plain, (~ (op(e1, e2) = op(e2, e1)) | spl30_98), inference(avatar_component_clause, [], [f834])).
fof(f7984, plain, (~ spl30_38 | ~ spl30_46), inference(avatar_split_clause, [], [f7983, f561, f527])).
fof(f7983, plain, (~ (e1 = op(e1, e2)) | ~ spl30_46), inference(forward_demodulation, [], [f194, f563])).
fof(f563, plain, ((e1 = op(e1, e0)) | ~ spl30_46), inference(avatar_component_clause, [], [f561])).
fof(f7982, plain, (~ spl30_40 | ~ spl30_56), inference(avatar_split_clause, [], [f7981, f603, f535])).
fof(f7981, plain, (~ (e3 = op(e1, e2)) | ~ spl30_56), inference(forward_demodulation, [], [f175, f605])).
fof(f175, plain, ~ (op(e0, e2) = op(e1, e2)), inference(cnf_transformation, [], [f3])).
fof(f7980, plain, (~ spl30_38 | ~ spl30_59 | spl30_139), inference(avatar_split_clause, [], [f7760, f1035, f616, f527])).
fof(f1035, plain, (spl30_139 <=> (e1 = op(e1, op(e0, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl30_139])])).
fof(f7760, plain, (~ (e1 = op(e1, e2)) | (~ spl30_59 | spl30_139)), inference(backward_demodulation, [], [f1037, f618])).
fof(f1037, plain, (~ (e1 = op(e1, op(e0, e1))) | spl30_139), inference(avatar_component_clause, [], [f1035])).
fof(f7979, plain, (~ spl30_40 | ~ spl30_44), inference(avatar_split_clause, [], [f7978, f552, f535])).
fof(f7978, plain, (~ (e3 = op(e1, e2)) | ~ spl30_44), inference(forward_demodulation, [], [f196, f554])).
fof(f554, plain, ((e3 = op(e1, e1)) | ~ spl30_44), inference(avatar_component_clause, [], [f552])).
fof(f7976, plain, (~ spl30_40 | spl30_212 | ~ spl30_214), inference(avatar_split_clause, [], [f7879, f1402, f1392, f535])).
fof(f1392, plain, (spl30_212 <=> (e3 = op(e1, op(e1, op(e1, e1))))), introduced(avatar_definition, [new_symbols(naming, [spl30_212])])).
fof(f1402, plain, (spl30_214 <=> (e2 = op(e1, op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl30_214])])).
fof(f7879, plain, (~ (e3 = op(e1, e2)) | (spl30_212 | ~ spl30_214)), inference(backward_demodulation, [], [f1394, f1403])).
fof(f1403, plain, ((e2 = op(e1, op(e1, e1))) | ~ spl30_214), inference(avatar_component_clause, [], [f1402])).
fof(f1394, plain, (~ (e3 = op(e1, op(e1, op(e1, e1)))) | spl30_212), inference(avatar_component_clause, [], [f1392])).
fof(f7975, plain, (~ spl30_37 | ~ spl30_214 | spl30_215), inference(avatar_split_clause, [], [f7880, f1407, f1402, f523])).
fof(f1407, plain, (spl30_215 <=> (e0 = op(e1, op(e1, op(e1, e1))))), introduced(avatar_definition, [new_symbols(naming, [spl30_215])])).
fof(f7880, plain, (~ (e0 = op(e1, e2)) | (~ spl30_214 | spl30_215)), inference(backward_demodulation, [], [f1409, f1403])).
fof(f1409, plain, (~ (e0 = op(e1, op(e1, op(e1, e1)))) | spl30_215), inference(avatar_component_clause, [], [f1407])).
fof(f7974, plain, (spl30_38 | ~ spl30_26 | ~ spl30_166), inference(avatar_split_clause, [], [f7952, f1166, f476, f527])).
fof(f1166, plain, (spl30_166 <=> (e1 = op(op(e2, e1), e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_166])])).
fof(f7952, plain, ((e1 = op(e1, e2)) | (~ spl30_26 | ~ spl30_166)), inference(backward_demodulation, [], [f1168, f478])).
fof(f1168, plain, ((e1 = op(op(e2, e1), e2)) | ~ spl30_166), inference(avatar_component_clause, [], [f1166])).
fof(f7972, plain, (~ spl30_36 | ~ spl30_44), inference(avatar_split_clause, [], [f7971, f552, f518])).
fof(f518, plain, (spl30_36 <=> (e3 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_36])])).
fof(f7971, plain, (~ (e3 = op(e1, e3)) | ~ spl30_44), inference(forward_demodulation, [], [f197, f554])).
fof(f7969, plain, (spl30_35 | ~ spl30_44 | ~ spl30_214), inference(avatar_split_clause, [], [f7936, f1402, f552, f514])).
fof(f7936, plain, ((e2 = op(e1, e3)) | (~ spl30_44 | ~ spl30_214)), inference(backward_demodulation, [], [f1403, f554])).
fof(f7967, plain, (~ spl30_24 | ~ spl30_32), inference(avatar_split_clause, [], [f7966, f501, f467])).
fof(f7966, plain, (~ (e3 = op(e2, e2)) | ~ spl30_32), inference(forward_demodulation, [], [f200, f503])).
fof(f200, plain, ~ (op(e2, e0) = op(e2, e2)), inference(cnf_transformation, [], [f3])).
fof(f7962, plain, (~ spl30_18 | ~ spl30_26), inference(avatar_split_clause, [], [f7961, f476, f442])).
fof(f442, plain, (spl30_18 <=> (e1 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_18])])).
fof(f7961, plain, (~ (e1 = op(e2, e3)) | ~ spl30_26), inference(forward_demodulation, [], [f203, f478])).
fof(f203, plain, ~ (op(e2, e1) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f7960, plain, (~ spl30_20 | ~ spl30_32), inference(avatar_split_clause, [], [f7959, f501, f450])).
fof(f7959, plain, (~ (e3 = op(e2, e3)) | ~ spl30_32), inference(forward_demodulation, [], [f201, f503])).
fof(f7957, plain, (~ spl30_26 | ~ spl30_37 | ~ spl30_166), inference(avatar_contradiction_clause, [], [f7956])).
fof(f7956, plain, ($false | (~ spl30_26 | ~ spl30_37 | ~ spl30_166)), inference(subsumption_resolution, [], [f7955, f211])).
fof(f211, plain, ~ (e0 = e1), inference(cnf_transformation, [], [f4])).
fof(f7955, plain, ((e0 = e1) | (~ spl30_26 | ~ spl30_37 | ~ spl30_166)), inference(forward_demodulation, [], [f7952, f525])).
fof(f7945, plain, (~ spl30_28 | ~ spl30_32), inference(avatar_split_clause, [], [f7941, f501, f484])).
fof(f484, plain, (spl30_28 <=> (e3 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_28])])).
fof(f7941, plain, (~ (e3 = op(e2, e1)) | ~ spl30_32), inference(backward_demodulation, [], [f199, f503])).
fof(f7939, plain, (~ spl30_4 | ~ spl30_44 | spl30_116), inference(avatar_contradiction_clause, [], [f7938])).
fof(f7938, plain, ($false | (~ spl30_4 | ~ spl30_44 | spl30_116)), inference(subsumption_resolution, [], [f7931, f384])).
fof(f7931, plain, (~ (e3 = op(e3, e3)) | (~ spl30_44 | spl30_116)), inference(backward_demodulation, [], [f924, f554])).
fof(f7937, plain, (~ spl30_28 | ~ spl30_44), inference(avatar_split_clause, [], [f7930, f552, f484])).
fof(f7930, plain, (~ (e3 = op(e2, e1)) | ~ spl30_44), inference(backward_demodulation, [], [f172, f554])).
fof(f7927, plain, (~ spl30_42 | ~ spl30_46), inference(avatar_split_clause, [], [f7916, f561, f544])).
fof(f7916, plain, (~ (e1 = op(e1, e1)) | ~ spl30_46), inference(backward_demodulation, [], [f193, f563])).
fof(f7926, plain, (~ spl30_30 | ~ spl30_46), inference(avatar_split_clause, [], [f7915, f561, f493])).
fof(f7915, plain, (~ (e1 = op(e2, e0)) | ~ spl30_46), inference(backward_demodulation, [], [f166, f563])).
fof(f7912, plain, (~ spl30_9 | spl30_31 | ~ spl30_73), inference(avatar_contradiction_clause, [], [f7911])).
fof(f7911, plain, ($false | (~ spl30_9 | spl30_31 | ~ spl30_73)), inference(subsumption_resolution, [], [f7910, f498])).
fof(f498, plain, (~ (e2 = op(e2, e0)) | spl30_31), inference(avatar_component_clause, [], [f497])).
fof(f7910, plain, ((e2 = op(e2, e0)) | (~ spl30_9 | ~ spl30_73)), inference(forward_demodulation, [], [f713, f406])).
fof(f713, plain, ((e2 = op(e2, op(e3, e1))) | ~ spl30_73), inference(avatar_component_clause, [], [f712])).
fof(f712, plain, (spl30_73 <=> (e2 = op(e2, op(e3, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl30_73])])).
fof(f7909, plain, (spl30_32 | ~ spl30_56 | ~ spl30_104), inference(avatar_split_clause, [], [f7908, f863, f603, f501])).
fof(f7908, plain, ((e3 = op(e2, e0)) | (~ spl30_56 | ~ spl30_104)), inference(forward_demodulation, [], [f865, f605])).
fof(f7907, plain, (~ spl30_35 | spl30_53 | ~ spl30_109), inference(avatar_contradiction_clause, [], [f7906])).
fof(f7906, plain, ($false | (~ spl30_35 | spl30_53 | ~ spl30_109)), inference(subsumption_resolution, [], [f7905, f592])).
fof(f592, plain, (~ (e0 = op(e0, e2)) | spl30_53), inference(avatar_component_clause, [], [f591])).
fof(f591, plain, (spl30_53 <=> (e0 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_53])])).
fof(f7905, plain, ((e0 = op(e0, e2)) | (~ spl30_35 | ~ spl30_109)), inference(forward_demodulation, [], [f888, f516])).
fof(f888, plain, ((e0 = op(e0, op(e1, e3))) | ~ spl30_109), inference(avatar_component_clause, [], [f887])).
fof(f887, plain, (spl30_109 <=> (e0 = op(e0, op(e1, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl30_109])])).
fof(f7901, plain, (~ spl30_42 | ~ spl30_50 | spl30_129), inference(avatar_split_clause, [], [f7900, f985, f578, f544])).
fof(f985, plain, (spl30_129 <=> (e1 = op(e1, op(e0, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl30_129])])).
fof(f7900, plain, (~ (e1 = op(e1, e1)) | (~ spl30_50 | spl30_129)), inference(forward_demodulation, [], [f987, f580])).
fof(f987, plain, (~ (e1 = op(e1, op(e0, e3))) | spl30_129), inference(avatar_component_clause, [], [f985])).
fof(f7894, plain, (spl30_32 | ~ spl30_17 | ~ spl30_161), inference(avatar_split_clause, [], [f7893, f1141, f438, f501])).
fof(f1141, plain, (spl30_161 <=> (e3 = op(e2, op(e2, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl30_161])])).
fof(f7893, plain, ((e3 = op(e2, e0)) | (~ spl30_17 | ~ spl30_161)), inference(forward_demodulation, [], [f1142, f440])).
fof(f1142, plain, ((e3 = op(e2, op(e2, e3))) | ~ spl30_161), inference(avatar_component_clause, [], [f1141])).
fof(f7891, plain, (~ spl30_28 | ~ spl30_35 | spl30_174), inference(avatar_split_clause, [], [f7890, f1204, f514, f484])).
fof(f1204, plain, (spl30_174 <=> (e3 = op(op(e1, e3), e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_174])])).
fof(f7890, plain, (~ (e3 = op(e2, e1)) | (~ spl30_35 | spl30_174)), inference(forward_demodulation, [], [f1205, f516])).
fof(f1205, plain, (~ (e3 = op(op(e1, e3), e1)) | spl30_174), inference(avatar_component_clause, [], [f1204])).
fof(f7873, plain, (~ spl30_46 | ~ spl30_37 | spl30_113), inference(avatar_split_clause, [], [f7592, f907, f523, f561])).
fof(f7592, plain, (~ (e1 = op(e1, e0)) | (~ spl30_37 | spl30_113)), inference(backward_demodulation, [], [f909, f525])).
fof(f7871, plain, (~ spl30_46 | ~ spl30_9 | spl30_74), inference(avatar_split_clause, [], [f7841, f717, f404, f561])).
fof(f717, plain, (spl30_74 <=> (e1 = op(e1, op(e3, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl30_74])])).
fof(f7841, plain, (~ (e1 = op(e1, e0)) | (~ spl30_9 | spl30_74)), inference(backward_demodulation, [], [f719, f406])).
fof(f719, plain, (~ (e1 = op(e1, op(e3, e1))) | spl30_74), inference(avatar_component_clause, [], [f717])).
fof(f7867, plain, (~ spl30_43 | ~ spl30_35), inference(avatar_split_clause, [], [f7866, f514, f548])).
fof(f548, plain, (spl30_43 <=> (e2 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_43])])).
fof(f7866, plain, (~ (e2 = op(e1, e1)) | ~ spl30_35), inference(forward_demodulation, [], [f197, f516])).
fof(f7865, plain, (~ spl30_43 | ~ spl30_59), inference(avatar_split_clause, [], [f7864, f616, f548])).
fof(f7864, plain, (~ (e2 = op(e1, e1)) | ~ spl30_59), inference(forward_demodulation, [], [f169, f618])).
fof(f169, plain, ~ (op(e0, e1) = op(e1, e1)), inference(cnf_transformation, [], [f3])).
fof(f7861, plain, (~ spl30_31 | ~ spl30_23), inference(avatar_split_clause, [], [f7860, f463, f497])).
fof(f7860, plain, (~ (e2 = op(e2, e0)) | ~ spl30_23), inference(forward_demodulation, [], [f200, f465])).
fof(f465, plain, ((e2 = op(e2, e2)) | ~ spl30_23), inference(avatar_component_clause, [], [f463])).
fof(f7858, plain, (~ spl30_32 | ~ spl30_15 | spl30_155), inference(avatar_split_clause, [], [f7828, f1112, f429, f501])).
fof(f1112, plain, (spl30_155 <=> (e3 = op(op(e3, e0), e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_155])])).
fof(f7828, plain, (~ (e3 = op(e2, e0)) | (~ spl30_15 | spl30_155)), inference(backward_demodulation, [], [f1114, f431])).
fof(f1114, plain, (~ (e3 = op(op(e3, e0), e0)) | spl30_155), inference(avatar_component_clause, [], [f1112])).
fof(f7831, plain, (~ spl30_3 | ~ spl30_15), inference(avatar_split_clause, [], [f7826, f429, f378])).
fof(f378, plain, (spl30_3 <=> (e2 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_3])])).
fof(f7826, plain, (~ (e2 = op(e3, e3)) | ~ spl30_15), inference(backward_demodulation, [], [f207, f431])).
fof(f7814, plain, (~ spl30_24 | ~ spl30_28), inference(avatar_split_clause, [], [f7805, f484, f467])).
fof(f7805, plain, (~ (e3 = op(e2, e2)) | ~ spl30_28), inference(backward_demodulation, [], [f202, f486])).
fof(f486, plain, ((e3 = op(e2, e1)) | ~ spl30_28), inference(avatar_component_clause, [], [f484])).
fof(f7813, plain, (~ spl30_12 | ~ spl30_28), inference(avatar_split_clause, [], [f7804, f484, f416])).
fof(f7804, plain, (~ (e3 = op(e3, e1)) | ~ spl30_28), inference(backward_demodulation, [], [f174, f486])).
fof(f174, plain, ~ (op(e2, e1) = op(e3, e1)), inference(cnf_transformation, [], [f3])).
fof(f7803, plain, (spl30_28 | ~ spl30_35 | ~ spl30_174), inference(avatar_split_clause, [], [f7797, f1204, f514, f484])).
fof(f7797, plain, ((e3 = op(e2, e1)) | (~ spl30_35 | ~ spl30_174)), inference(backward_demodulation, [], [f1206, f516])).
fof(f1206, plain, ((e3 = op(op(e1, e3), e1)) | ~ spl30_174), inference(avatar_component_clause, [], [f1204])).
fof(f7798, plain, (~ spl30_3 | ~ spl30_35), inference(avatar_split_clause, [], [f7790, f514, f378])).
fof(f7790, plain, (~ (e2 = op(e3, e3)) | ~ spl30_35), inference(backward_demodulation, [], [f185, f516])).
fof(f7786, plain, (~ spl30_4 | ~ spl30_56 | spl30_132), inference(avatar_split_clause, [], [f7778, f1000, f603, f382])).
fof(f1000, plain, (spl30_132 <=> (e3 = op(e3, op(e0, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl30_132])])).
fof(f7778, plain, (~ (e3 = op(e3, e3)) | (~ spl30_56 | spl30_132)), inference(backward_demodulation, [], [f1002, f605])).
fof(f1002, plain, (~ (e3 = op(e3, op(e0, e2))) | spl30_132), inference(avatar_component_clause, [], [f1000])).
fof(f7785, plain, (~ spl30_24 | ~ spl30_56), inference(avatar_split_clause, [], [f7777, f603, f467])).
fof(f7777, plain, (~ (e3 = op(e2, e2)) | ~ spl30_56), inference(backward_demodulation, [], [f176, f605])).
fof(f7774, plain, (~ spl30_35 | ~ spl30_59 | ~ spl30_125), inference(avatar_split_clause, [], [f7766, f966, f616, f514])).
fof(f7766, plain, (~ (e2 = op(e1, e3)) | (~ spl30_59 | ~ spl30_125)), inference(backward_demodulation, [], [f7714, f618])).
fof(f7714, plain, (~ (op(e0, e1) = op(e1, e3)) | ~ spl30_125), inference(backward_demodulation, [], [f195, f968])).
fof(f195, plain, ~ (op(e1, e0) = op(e1, e3)), inference(cnf_transformation, [], [f3])).
fof(f7773, plain, (~ spl30_15 | ~ spl30_59 | ~ spl30_125), inference(avatar_split_clause, [], [f7765, f966, f616, f429])).
fof(f7765, plain, (~ (e2 = op(e3, e0)) | (~ spl30_59 | ~ spl30_125)), inference(backward_demodulation, [], [f7713, f618])).
fof(f7713, plain, (~ (op(e0, e1) = op(e3, e0)) | ~ spl30_125), inference(backward_demodulation, [], [f167, f968])).
fof(f7770, plain, (spl30_47 | ~ spl30_59 | ~ spl30_125), inference(avatar_split_clause, [], [f7758, f966, f616, f565])).
fof(f7758, plain, ((e2 = op(e1, e0)) | (~ spl30_59 | ~ spl30_125)), inference(backward_demodulation, [], [f968, f618])).
fof(f7769, plain, (~ spl30_55 | ~ spl30_59), inference(avatar_split_clause, [], [f7756, f616, f599])).
fof(f599, plain, (spl30_55 <=> (e2 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_55])])).
fof(f7756, plain, (~ (e2 = op(e0, e2)) | ~ spl30_59), inference(backward_demodulation, [], [f190, f618])).
fof(f7768, plain, (~ spl30_27 | ~ spl30_59), inference(avatar_split_clause, [], [f7754, f616, f480])).
fof(f7754, plain, (~ (e2 = op(e2, e1)) | ~ spl30_59), inference(backward_demodulation, [], [f170, f618])).
fof(f170, plain, ~ (op(e0, e1) = op(e2, e1)), inference(cnf_transformation, [], [f3])).
fof(f7753, plain, (~ spl30_57 | ~ spl30_61), inference(avatar_split_clause, [], [f7745, f625, f608])).
fof(f7745, plain, (~ (e0 = op(e0, e1)) | ~ spl30_61), inference(backward_demodulation, [], [f187, f627])).
fof(f187, plain, ~ (op(e0, e0) = op(e0, e1)), inference(cnf_transformation, [], [f3])).
fof(f7752, plain, (~ spl30_13 | ~ spl30_61), inference(avatar_split_clause, [], [f7744, f625, f421])).
fof(f421, plain, (spl30_13 <=> (e0 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_13])])).
fof(f7744, plain, (~ (e0 = op(e3, e0)) | ~ spl30_61), inference(backward_demodulation, [], [f165, f627])).
fof(f165, plain, ~ (op(e0, e0) = op(e3, e0)), inference(cnf_transformation, [], [f3])).
fof(f7742, plain, (~ spl30_12 | ~ spl30_6 | spl30_66), inference(avatar_split_clause, [], [f7741, f678, f391, f416])).
fof(f678, plain, (spl30_66 <=> (e3 = op(e3, op(e3, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl30_66])])).
fof(f7741, plain, (~ (e3 = op(e3, e1)) | (~ spl30_6 | spl30_66)), inference(forward_demodulation, [], [f680, f393])).
fof(f680, plain, (~ (e3 = op(e3, op(e3, e2))) | spl30_66), inference(avatar_component_clause, [], [f678])).
fof(f7738, plain, (~ spl30_24 | ~ spl30_90 | spl30_218), inference(avatar_split_clause, [], [f7735, f1422, f795, f467])).
fof(f795, plain, (spl30_90 <=> (e2 = op(e2, op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl30_90])])).
fof(f7735, plain, (~ (e3 = op(e2, e2)) | (~ spl30_90 | spl30_218)), inference(backward_demodulation, [], [f1424, f796])).
fof(f796, plain, ((e2 = op(e2, op(e2, e2))) | ~ spl30_90), inference(avatar_component_clause, [], [f795])).
fof(f7724, plain, (~ spl30_137 | spl30_121 | ~ spl30_125), inference(avatar_split_clause, [], [f7723, f966, f946, f1025])).
fof(f946, plain, (spl30_121 <=> (e3 = op(e3, op(e1, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl30_121])])).
fof(f7723, plain, (~ (e3 = op(e3, op(e0, e1))) | (spl30_121 | ~ spl30_125)), inference(forward_demodulation, [], [f948, f968])).
fof(f948, plain, (~ (e3 = op(e3, op(e1, e0))) | spl30_121), inference(avatar_component_clause, [], [f946])).
fof(f7705, plain, (spl30_56 | ~ spl30_17 | ~ spl30_160), inference(avatar_split_clause, [], [f7704, f1136, f438, f603])).
fof(f1136, plain, (spl30_160 <=> (e3 = op(op(e2, e3), e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_160])])).
fof(f7704, plain, ((e3 = op(e0, e2)) | (~ spl30_17 | ~ spl30_160)), inference(forward_demodulation, [], [f1138, f440])).
fof(f1138, plain, ((e3 = op(op(e2, e3), e2)) | ~ spl30_160), inference(avatar_component_clause, [], [f1136])).
fof(f7701, plain, (~ spl30_47 | ~ spl30_37 | spl30_179), inference(avatar_split_clause, [], [f7700, f1228, f523, f565])).
fof(f1228, plain, (spl30_179 <=> (e2 = op(e1, op(e1, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl30_179])])).
fof(f7700, plain, (~ (e2 = op(e1, e0)) | (~ spl30_37 | spl30_179)), inference(forward_demodulation, [], [f1230, f525])).
fof(f1230, plain, (~ (e2 = op(e1, op(e1, e2))) | spl30_179), inference(avatar_component_clause, [], [f1228])).
fof(f7687, plain, (spl30_59 | ~ spl30_37 | ~ spl30_178), inference(avatar_split_clause, [], [f7595, f1223, f523, f616])).
fof(f1223, plain, (spl30_178 <=> (e2 = op(op(e1, e2), e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_178])])).
fof(f7595, plain, ((e2 = op(e0, e1)) | (~ spl30_37 | ~ spl30_178)), inference(backward_demodulation, [], [f1225, f525])).
fof(f1225, plain, ((e2 = op(op(e1, e2), e1)) | ~ spl30_178), inference(avatar_component_clause, [], [f1223])).
fof(f7655, plain, (~ spl30_6 | spl30_35 | ~ spl30_148), inference(avatar_contradiction_clause, [], [f7654])).
fof(f7654, plain, ($false | (~ spl30_6 | spl30_35 | ~ spl30_148)), inference(subsumption_resolution, [], [f7651, f515])).
fof(f515, plain, (~ (e2 = op(e1, e3)) | spl30_35), inference(avatar_component_clause, [], [f514])).
fof(f7651, plain, ((e2 = op(e1, e3)) | (~ spl30_6 | ~ spl30_148)), inference(backward_demodulation, [], [f1081, f393])).
fof(f1081, plain, ((e2 = op(op(e3, e2), e3)) | ~ spl30_148), inference(avatar_component_clause, [], [f1079])).
fof(f1079, plain, (spl30_148 <=> (e2 = op(op(e3, e2), e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_148])])).
fof(f7640, plain, (~ spl30_1 | ~ spl30_13), inference(avatar_split_clause, [], [f7631, f421, f370])).
fof(f7631, plain, (~ (e0 = op(e3, e3)) | ~ spl30_13), inference(backward_demodulation, [], [f207, f423])).
fof(f423, plain, ((e0 = op(e3, e0)) | ~ spl30_13), inference(avatar_component_clause, [], [f421])).
fof(f7629, plain, (~ spl30_17 | ~ spl30_24 | ~ spl30_30 | spl30_219), inference(avatar_contradiction_clause, [], [f7628])).
fof(f7628, plain, ($false | (~ spl30_17 | ~ spl30_24 | ~ spl30_30 | spl30_219)), inference(subsumption_resolution, [], [f7622, f495])).
fof(f7622, plain, (~ (e1 = op(e2, e0)) | (~ spl30_17 | ~ spl30_24 | spl30_219)), inference(backward_demodulation, [], [f7506, f440])).
fof(f7506, plain, (~ (e1 = op(e2, op(e2, e3))) | (~ spl30_24 | spl30_219)), inference(forward_demodulation, [], [f1429, f469])).
fof(f1429, plain, (~ (e1 = op(e2, op(e2, op(e2, e2)))) | spl30_219), inference(avatar_component_clause, [], [f1427])).
fof(f1427, plain, (spl30_219 <=> (e1 = op(e2, op(e2, op(e2, e2))))), introduced(avatar_definition, [new_symbols(naming, [spl30_219])])).
fof(f7610, plain, (~ spl30_18 | ~ spl30_30), inference(avatar_split_clause, [], [f7603, f493, f442])).
fof(f7603, plain, (~ (e1 = op(e2, e3)) | ~ spl30_30), inference(backward_demodulation, [], [f201, f495])).
fof(f7601, plain, (~ spl30_37 | ~ spl30_57 | ~ spl30_178), inference(avatar_contradiction_clause, [], [f7600])).
fof(f7600, plain, ($false | (~ spl30_37 | ~ spl30_57 | ~ spl30_178)), inference(subsumption_resolution, [], [f7599, f212])).
fof(f212, plain, ~ (e0 = e2), inference(cnf_transformation, [], [f4])).
fof(f7599, plain, ((e0 = e2) | (~ spl30_37 | ~ spl30_57 | ~ spl30_178)), inference(forward_demodulation, [], [f7595, f610])).
fof(f610, plain, ((e0 = op(e0, e1)) | ~ spl30_57), inference(avatar_component_clause, [], [f608])).
fof(f7584, plain, (~ spl30_38 | ~ spl30_42), inference(avatar_split_clause, [], [f7576, f544, f527])).
fof(f7576, plain, (~ (e1 = op(e1, e2)) | ~ spl30_42), inference(backward_demodulation, [], [f196, f546])).
fof(f7570, plain, (~ spl30_18 | ~ spl30_50), inference(avatar_split_clause, [], [f7560, f578, f442])).
fof(f7560, plain, (~ (e1 = op(e2, e3)) | ~ spl30_50), inference(backward_demodulation, [], [f182, f580])).
fof(f182, plain, ~ (op(e0, e3) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f7557, plain, (~ spl30_51 | ~ spl30_55), inference(avatar_split_clause, [], [f7546, f599, f582])).
fof(f7546, plain, (~ (e2 = op(e0, e3)) | ~ spl30_55), inference(backward_demodulation, [], [f192, f601])).
fof(f601, plain, ((e2 = op(e0, e2)) | ~ spl30_55), inference(avatar_component_clause, [], [f599])).
fof(f7556, plain, (~ spl30_7 | ~ spl30_55), inference(avatar_split_clause, [], [f7545, f599, f395])).
fof(f7545, plain, (~ (e2 = op(e3, e2)) | ~ spl30_55), inference(backward_demodulation, [], [f177, f601])).
fof(f7541, plain, (~ spl30_53 | ~ spl30_57), inference(avatar_split_clause, [], [f7530, f608, f591])).
fof(f7530, plain, (~ (e0 = op(e0, e2)) | ~ spl30_57), inference(backward_demodulation, [], [f190, f610])).
fof(f7540, plain, (~ spl30_41 | ~ spl30_57), inference(avatar_split_clause, [], [f7529, f608, f540])).
fof(f7529, plain, (~ (e0 = op(e1, e1)) | ~ spl30_57), inference(backward_demodulation, [], [f169, f610])).
fof(f7511, plain, (spl30_37 | ~ spl30_47 | ~ spl30_185), inference(avatar_split_clause, [], [f7510, f1258, f565, f523])).
fof(f7510, plain, ((e0 = op(e1, e2)) | (~ spl30_47 | ~ spl30_185)), inference(forward_demodulation, [], [f1259, f567])).
fof(f1259, plain, ((e0 = op(e1, op(e1, e0))) | ~ spl30_185), inference(avatar_component_clause, [], [f1258])).
fof(f7509, plain, (spl30_50 | ~ spl30_64 | ~ spl30_205), inference(avatar_split_clause, [], [f7508, f1358, f637, f578])).
fof(f637, plain, (spl30_64 <=> (op(e0, e0) = e3)), introduced(avatar_definition, [new_symbols(naming, [spl30_64])])).
fof(f1358, plain, (spl30_205 <=> (e1 = op(e0, op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl30_205])])).
fof(f7508, plain, ((e1 = op(e0, e3)) | (~ spl30_64 | ~ spl30_205)), inference(forward_demodulation, [], [f1359, f639])).
fof(f639, plain, ((op(e0, e0) = e3) | ~ spl30_64), inference(avatar_component_clause, [], [f637])).
fof(f1359, plain, ((e1 = op(e0, op(e0, e0))) | ~ spl30_205), inference(avatar_component_clause, [], [f1358])).
fof(f7496, plain, (~ spl30_52 | ~ spl30_64 | spl30_210), inference(avatar_split_clause, [], [f7297, f1382, f637, f586])).
fof(f586, plain, (spl30_52 <=> (e3 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_52])])).
fof(f1382, plain, (spl30_210 <=> (e3 = op(e0, op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl30_210])])).
fof(f7297, plain, (~ (e3 = op(e0, e3)) | (~ spl30_64 | spl30_210)), inference(forward_demodulation, [], [f1384, f639])).
fof(f1384, plain, (~ (e3 = op(e0, op(e0, e0))) | spl30_210), inference(avatar_component_clause, [], [f1382])).
fof(f7485, plain, (~ spl30_32 | ~ spl30_64), inference(avatar_split_clause, [], [f7484, f637, f501])).
fof(f7484, plain, (~ (e3 = op(e2, e0)) | ~ spl30_64), inference(forward_demodulation, [], [f164, f639])).
fof(f164, plain, ~ (op(e0, e0) = op(e2, e0)), inference(cnf_transformation, [], [f3])).
fof(f7481, plain, (~ spl30_19 | ~ spl30_27), inference(avatar_split_clause, [], [f7480, f480, f446])).
fof(f7480, plain, (~ (e2 = op(e2, e3)) | ~ spl30_27), inference(forward_demodulation, [], [f203, f482])).
fof(f482, plain, ((e2 = op(e2, e1)) | ~ spl30_27), inference(avatar_component_clause, [], [f480])).
fof(f7479, plain, (~ spl30_20 | ~ spl30_36), inference(avatar_split_clause, [], [f7478, f518, f450])).
fof(f7478, plain, (~ (e3 = op(e2, e3)) | ~ spl30_36), inference(forward_demodulation, [], [f184, f520])).
fof(f520, plain, ((e3 = op(e1, e3)) | ~ spl30_36), inference(avatar_component_clause, [], [f518])).
fof(f7475, plain, (~ spl30_17 | ~ spl30_24 | spl30_217), inference(avatar_split_clause, [], [f7451, f1418, f467, f438])).
fof(f7451, plain, (~ (e0 = op(e2, e3)) | (~ spl30_24 | spl30_217)), inference(backward_demodulation, [], [f1420, f469])).
fof(f1420, plain, (~ (e0 = op(e2, op(e2, e2))) | spl30_217), inference(avatar_component_clause, [], [f1418])).
fof(f7455, plain, (~ spl30_24 | spl30_49 | ~ spl30_92), inference(avatar_contradiction_clause, [], [f7454])).
fof(f7454, plain, ($false | (~ spl30_24 | spl30_49 | ~ spl30_92)), inference(subsumption_resolution, [], [f7450, f575])).
fof(f7450, plain, ((e0 = op(e0, e3)) | (~ spl30_24 | ~ spl30_92)), inference(backward_demodulation, [], [f806, f469])).
fof(f806, plain, ((e0 = op(e0, op(e2, e2))) | ~ spl30_92), inference(avatar_component_clause, [], [f805])).
fof(f805, plain, (spl30_92 <=> (e0 = op(e0, op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl30_92])])).
fof(f7453, plain, (~ spl30_4 | ~ spl30_24 | spl30_89), inference(avatar_split_clause, [], [f7448, f790, f467, f382])).
fof(f790, plain, (spl30_89 <=> (e3 = op(e3, op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl30_89])])).
fof(f7448, plain, (~ (e3 = op(e3, e3)) | (~ spl30_24 | spl30_89)), inference(backward_demodulation, [], [f792, f469])).
fof(f792, plain, (~ (e3 = op(e3, op(e2, e2))) | spl30_89), inference(avatar_component_clause, [], [f790])).
fof(f7446, plain, (spl30_22 | ~ spl30_27 | ~ spl30_167), inference(avatar_contradiction_clause, [], [f7445])).
fof(f7445, plain, ($false | (spl30_22 | ~ spl30_27 | ~ spl30_167)), inference(subsumption_resolution, [], [f7437, f460])).
fof(f460, plain, (~ (e1 = op(e2, e2)) | spl30_22), inference(avatar_component_clause, [], [f459])).
fof(f7437, plain, ((e1 = op(e2, e2)) | (~ spl30_27 | ~ spl30_167)), inference(backward_demodulation, [], [f1172, f482])).
fof(f1172, plain, ((e1 = op(e2, op(e2, e1))) | ~ spl30_167), inference(avatar_component_clause, [], [f1171])).
fof(f1171, plain, (spl30_167 <=> (e1 = op(e2, op(e2, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl30_167])])).
fof(f7438, plain, (~ spl30_23 | ~ spl30_27), inference(avatar_split_clause, [], [f7430, f480, f463])).
fof(f7430, plain, (~ (e2 = op(e2, e2)) | ~ spl30_27), inference(backward_demodulation, [], [f202, f482])).
fof(f7424, plain, (~ spl30_4 | ~ spl30_36), inference(avatar_split_clause, [], [f7412, f518, f382])).
fof(f7412, plain, (~ (e3 = op(e3, e3)) | ~ spl30_36), inference(backward_demodulation, [], [f185, f520])).
fof(f7375, plain, (~ spl30_53 | spl30_63 | ~ spl30_193), inference(avatar_contradiction_clause, [], [f7374])).
fof(f7374, plain, ($false | (~ spl30_53 | spl30_63 | ~ spl30_193)), inference(subsumption_resolution, [], [f7368, f634])).
fof(f634, plain, (~ (op(e0, e0) = e2) | spl30_63), inference(avatar_component_clause, [], [f633])).
fof(f633, plain, (spl30_63 <=> (op(e0, e0) = e2)), introduced(avatar_definition, [new_symbols(naming, [spl30_63])])).
fof(f7368, plain, ((op(e0, e0) = e2) | (~ spl30_53 | ~ spl30_193)), inference(backward_demodulation, [], [f1297, f593])).
fof(f593, plain, ((e0 = op(e0, e2)) | ~ spl30_53), inference(avatar_component_clause, [], [f591])).
fof(f1297, plain, ((e2 = op(e0, op(e0, e2))) | ~ spl30_193), inference(avatar_component_clause, [], [f1296])).
fof(f1296, plain, (spl30_193 <=> (e2 = op(e0, op(e0, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl30_193])])).
fof(f7369, plain, (~ spl30_5 | ~ spl30_53), inference(avatar_split_clause, [], [f7358, f591, f387])).
fof(f7358, plain, (~ (e0 = op(e3, e2)) | ~ spl30_53), inference(backward_demodulation, [], [f177, f593])).
fof(f7355, plain, (spl30_41 | ~ spl30_58 | ~ spl30_195), inference(avatar_split_clause, [], [f7347, f1305, f612, f540])).
fof(f612, plain, (spl30_58 <=> (e1 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_58])])).
fof(f7347, plain, ((e0 = op(e1, e1)) | (~ spl30_58 | ~ spl30_195)), inference(backward_demodulation, [], [f1306, f614])).
fof(f614, plain, ((e1 = op(e0, e1)) | ~ spl30_58), inference(avatar_component_clause, [], [f612])).
fof(f7351, plain, (~ spl30_54 | ~ spl30_58), inference(avatar_split_clause, [], [f7341, f612, f595])).
fof(f7341, plain, (~ (e1 = op(e0, e2)) | ~ spl30_58), inference(backward_demodulation, [], [f190, f614])).
fof(f7350, plain, (~ spl30_42 | ~ spl30_58), inference(avatar_split_clause, [], [f7338, f612, f544])).
fof(f7338, plain, (~ (e1 = op(e1, e1)) | ~ spl30_58), inference(backward_demodulation, [], [f169, f614])).
fof(f7321, plain, (~ spl30_57 | ~ spl30_18 | spl30_87), inference(avatar_split_clause, [], [f7320, f780, f442, f608])).
fof(f780, plain, (spl30_87 <=> (e0 = op(e0, op(e2, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl30_87])])).
fof(f7320, plain, (~ (e0 = op(e0, e1)) | (~ spl30_18 | spl30_87)), inference(forward_demodulation, [], [f782, f444])).
fof(f444, plain, ((e1 = op(e2, e3)) | ~ spl30_18), inference(avatar_component_clause, [], [f442])).
fof(f782, plain, (~ (e0 = op(e0, op(e2, e3))) | spl30_87), inference(avatar_component_clause, [], [f780])).
fof(f7318, plain, (~ spl30_23 | ~ spl30_47 | spl30_122), inference(avatar_split_clause, [], [f7317, f951, f565, f463])).
fof(f951, plain, (spl30_122 <=> (e2 = op(e2, op(e1, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl30_122])])).
fof(f7317, plain, (~ (e2 = op(e2, e2)) | (~ spl30_47 | spl30_122)), inference(forward_demodulation, [], [f953, f567])).
fof(f953, plain, (~ (e2 = op(e2, op(e1, e0))) | spl30_122), inference(avatar_component_clause, [], [f951])).
fof(f7307, plain, (~ spl30_28 | ~ spl30_18 | spl30_161), inference(avatar_split_clause, [], [f7306, f1141, f442, f484])).
fof(f7306, plain, (~ (e3 = op(e2, e1)) | (~ spl30_18 | spl30_161)), inference(forward_demodulation, [], [f1143, f444])).
fof(f1143, plain, (~ (e3 = op(e2, op(e2, e3))) | spl30_161), inference(avatar_component_clause, [], [f1141])).
fof(f7304, plain, (~ spl30_29 | ~ spl30_51 | ~ spl30_188), inference(avatar_contradiction_clause, [], [f7303])).
fof(f7303, plain, ($false | (~ spl30_29 | ~ spl30_51 | ~ spl30_188)), inference(subsumption_resolution, [], [f7302, f213])).
fof(f213, plain, ~ (e0 = e3), inference(cnf_transformation, [], [f4])).
fof(f7302, plain, ((e0 = e3) | (~ spl30_29 | ~ spl30_51 | ~ spl30_188)), inference(forward_demodulation, [], [f7301, f491])).
fof(f491, plain, ((e0 = op(e2, e0)) | ~ spl30_29), inference(avatar_component_clause, [], [f489])).
fof(f489, plain, (spl30_29 <=> (e0 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_29])])).
fof(f7301, plain, ((e3 = op(e2, e0)) | (~ spl30_51 | ~ spl30_188)), inference(forward_demodulation, [], [f1274, f584])).
fof(f7292, plain, (~ spl30_28 | spl30_218 | ~ spl30_220), inference(avatar_split_clause, [], [f7290, f1432, f1422, f484])).
fof(f1432, plain, (spl30_220 <=> (e1 = op(e2, op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl30_220])])).
fof(f7290, plain, (~ (e3 = op(e2, e1)) | (spl30_218 | ~ spl30_220)), inference(backward_demodulation, [], [f1424, f1433])).
fof(f1433, plain, ((e1 = op(e2, op(e2, e2))) | ~ spl30_220), inference(avatar_component_clause, [], [f1432])).
fof(f7285, plain, (~ spl30_54 | ~ spl30_51 | ~ spl30_64 | spl30_209), inference(avatar_split_clause, [], [f7284, f1377, f637, f582, f595])).
fof(f1377, plain, (spl30_209 <=> (e1 = op(e0, op(e0, op(e0, e0))))), introduced(avatar_definition, [new_symbols(naming, [spl30_209])])).
fof(f7284, plain, (~ (e1 = op(e0, e2)) | (~ spl30_51 | ~ spl30_64 | spl30_209)), inference(forward_demodulation, [], [f7283, f584])).
fof(f7283, plain, (~ (e1 = op(e0, op(e0, e3))) | (~ spl30_64 | spl30_209)), inference(forward_demodulation, [], [f1379, f639])).
fof(f1379, plain, (~ (e1 = op(e0, op(e0, op(e0, e0)))) | spl30_209), inference(avatar_component_clause, [], [f1377])).
fof(f7276, plain, (~ spl30_56 | ~ spl30_51 | spl30_189), inference(avatar_split_clause, [], [f7165, f1277, f582, f603])).
fof(f1277, plain, (spl30_189 <=> (e3 = op(e0, op(e0, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl30_189])])).
fof(f7165, plain, (~ (e3 = op(e0, e2)) | (~ spl30_51 | spl30_189)), inference(backward_demodulation, [], [f1279, f584])).
fof(f1279, plain, (~ (e3 = op(e0, op(e0, e3))) | spl30_189), inference(avatar_component_clause, [], [f1277])).
fof(f7273, plain, (~ spl30_53 | ~ spl30_51 | spl30_130), inference(avatar_split_clause, [], [f7162, f990, f582, f591])).
fof(f990, plain, (spl30_130 <=> (e0 = op(e0, op(e0, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl30_130])])).
fof(f7162, plain, (~ (e0 = op(e0, e2)) | (~ spl30_51 | spl30_130)), inference(backward_demodulation, [], [f992, f584])).
fof(f992, plain, (~ (e0 = op(e0, op(e0, e3))) | spl30_130), inference(avatar_component_clause, [], [f990])).
fof(f7247, plain, (~ spl30_6 | ~ spl30_11 | spl30_153), inference(avatar_split_clause, [], [f7241, f1103, f412, f391])).
fof(f1103, plain, (spl30_153 <=> (e1 = op(e3, op(e3, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl30_153])])).
fof(f7241, plain, (~ (e1 = op(e3, e2)) | (~ spl30_11 | spl30_153)), inference(backward_demodulation, [], [f1105, f414])).
fof(f414, plain, ((e2 = op(e3, e1)) | ~ spl30_11), inference(avatar_component_clause, [], [f412])).
fof(f1105, plain, (~ (e1 = op(e3, op(e3, e1))) | spl30_153), inference(avatar_component_clause, [], [f1103])).
fof(f7246, plain, (~ spl30_11 | ~ spl30_18 | spl30_152), inference(avatar_contradiction_clause, [], [f7245])).
fof(f7245, plain, ($false | (~ spl30_11 | ~ spl30_18 | spl30_152)), inference(subsumption_resolution, [], [f7240, f444])).
fof(f7240, plain, (~ (e1 = op(e2, e3)) | (~ spl30_11 | spl30_152)), inference(backward_demodulation, [], [f1099, f414])).
fof(f1099, plain, (~ (e1 = op(op(e3, e1), e3)) | spl30_152), inference(avatar_component_clause, [], [f1098])).
fof(f1098, plain, (spl30_152 <=> (e1 = op(op(e3, e1), e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_152])])).
fof(f7242, plain, (~ spl30_3 | ~ spl30_11), inference(avatar_split_clause, [], [f7233, f412, f378])).
fof(f7233, plain, (~ (e2 = op(e3, e3)) | ~ spl30_11), inference(backward_demodulation, [], [f209, f414])).
fof(f209, plain, ~ (op(e3, e1) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f7217, plain, (~ spl30_23 | spl30_53 | ~ spl30_92), inference(avatar_contradiction_clause, [], [f7216])).
fof(f7216, plain, ($false | (~ spl30_23 | spl30_53 | ~ spl30_92)), inference(subsumption_resolution, [], [f7213, f592])).
fof(f7213, plain, ((e0 = op(e0, e2)) | (~ spl30_23 | ~ spl30_92)), inference(backward_demodulation, [], [f806, f465])).
fof(f7203, plain, (~ spl30_9 | ~ spl30_33 | spl30_76), inference(avatar_split_clause, [], [f7196, f727, f506, f404])).
fof(f7196, plain, (~ (e0 = op(e3, e1)) | (~ spl30_33 | spl30_76)), inference(backward_demodulation, [], [f728, f508])).
fof(f728, plain, (~ (op(e1, e3) = op(e3, e1)) | spl30_76), inference(avatar_component_clause, [], [f727])).
fof(f7191, plain, (spl30_11 | ~ spl30_40 | ~ spl30_178), inference(avatar_split_clause, [], [f7182, f1223, f535, f412])).
fof(f7182, plain, ((e2 = op(e3, e1)) | (~ spl30_40 | ~ spl30_178)), inference(backward_demodulation, [], [f1225, f537])).
fof(f7190, plain, (~ spl30_6 | ~ spl30_40 | spl30_177), inference(avatar_split_clause, [], [f7181, f1218, f535, f391])).
fof(f1218, plain, (spl30_177 <=> (e1 = op(op(e1, e2), e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_177])])).
fof(f7181, plain, (~ (e1 = op(e3, e2)) | (~ spl30_40 | spl30_177)), inference(backward_demodulation, [], [f1220, f537])).
fof(f1220, plain, (~ (e1 = op(op(e1, e2), e2)) | spl30_177), inference(avatar_component_clause, [], [f1218])).
fof(f7188, plain, (spl30_28 | ~ spl30_40 | ~ spl30_98), inference(avatar_split_clause, [], [f7177, f834, f535, f484])).
fof(f7177, plain, ((e3 = op(e2, e1)) | (~ spl30_40 | ~ spl30_98)), inference(backward_demodulation, [], [f836, f537])).
fof(f7187, plain, (~ spl30_36 | ~ spl30_40), inference(avatar_split_clause, [], [f7176, f535, f518])).
fof(f7176, plain, (~ (e3 = op(e1, e3)) | ~ spl30_40), inference(backward_demodulation, [], [f198, f537])).
fof(f7171, plain, (spl30_17 | ~ spl30_51 | ~ spl30_187), inference(avatar_contradiction_clause, [], [f7170])).
fof(f7170, plain, ($false | (spl30_17 | ~ spl30_51 | ~ spl30_187)), inference(subsumption_resolution, [], [f7163, f439])).
fof(f439, plain, (~ (e0 = op(e2, e3)) | spl30_17), inference(avatar_component_clause, [], [f438])).
fof(f7163, plain, ((e0 = op(e2, e3)) | (~ spl30_51 | ~ spl30_187)), inference(backward_demodulation, [], [f1268, f584])).
fof(f7167, plain, (~ spl30_23 | ~ spl30_51 | spl30_128), inference(avatar_split_clause, [], [f7160, f980, f582, f463])).
fof(f980, plain, (spl30_128 <=> (e2 = op(e2, op(e0, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl30_128])])).
fof(f7160, plain, (~ (e2 = op(e2, e2)) | (~ spl30_51 | spl30_128)), inference(backward_demodulation, [], [f982, f584])).
fof(f982, plain, (~ (e2 = op(e2, op(e0, e3))) | spl30_128), inference(avatar_component_clause, [], [f980])).
fof(f7166, plain, (~ spl30_3 | ~ spl30_51), inference(avatar_split_clause, [], [f7157, f582, f378])).
fof(f7157, plain, (~ (e2 = op(e3, e3)) | ~ spl30_51), inference(backward_demodulation, [], [f183, f584])).
fof(f183, plain, ~ (op(e0, e3) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f7154, plain, (~ spl30_27 | ~ spl30_54 | spl30_133), inference(avatar_split_clause, [], [f7148, f1005, f595, f480])).
fof(f1005, plain, (spl30_133 <=> (e2 = op(e2, op(e0, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl30_133])])).
fof(f7148, plain, (~ (e2 = op(e2, e1)) | (~ spl30_54 | spl30_133)), inference(backward_demodulation, [], [f1007, f597])).
fof(f1007, plain, (~ (e2 = op(e2, op(e0, e2))) | spl30_133), inference(avatar_component_clause, [], [f1005])).
fof(f7153, plain, (~ spl30_6 | ~ spl30_54), inference(avatar_split_clause, [], [f7144, f595, f391])).
fof(f7144, plain, (~ (e1 = op(e3, e2)) | ~ spl30_54), inference(backward_demodulation, [], [f177, f597])).
fof(f7138, plain, (~ spl30_16 | ~ spl30_57 | spl30_137), inference(avatar_split_clause, [], [f7130, f1025, f608, f433])).
fof(f7130, plain, (~ (e3 = op(e3, e0)) | (~ spl30_57 | spl30_137)), inference(backward_demodulation, [], [f1027, f610])).
fof(f7137, plain, (~ spl30_49 | ~ spl30_57), inference(avatar_split_clause, [], [f7126, f608, f574])).
fof(f7126, plain, (~ (e0 = op(e0, e3)) | ~ spl30_57), inference(backward_demodulation, [], [f191, f610])).
fof(f191, plain, ~ (op(e0, e1) = op(e0, e3)), inference(cnf_transformation, [], [f3])).
fof(f7136, plain, (~ spl30_9 | ~ spl30_57), inference(avatar_split_clause, [], [f7124, f608, f404])).
fof(f7124, plain, (~ (e0 = op(e3, e1)) | ~ spl30_57), inference(backward_demodulation, [], [f171, f610])).
fof(f7120, plain, (~ spl30_51 | ~ spl30_64 | spl30_208), inference(avatar_split_clause, [], [f7113, f1372, f637, f582])).
fof(f1372, plain, (spl30_208 <=> (e2 = op(e0, op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl30_208])])).
fof(f7113, plain, (~ (e2 = op(e0, e3)) | (~ spl30_64 | spl30_208)), inference(backward_demodulation, [], [f1374, f639])).
fof(f1374, plain, (~ (e2 = op(e0, op(e0, e0))) | spl30_208), inference(avatar_component_clause, [], [f1372])).
fof(f7117, plain, (~ spl30_4 | ~ spl30_64 | spl30_142), inference(avatar_split_clause, [], [f7108, f1050, f637, f382])).
fof(f7108, plain, (~ (e3 = op(e3, e3)) | (~ spl30_64 | spl30_142)), inference(backward_demodulation, [], [f1052, f639])).
fof(f7116, plain, (~ spl30_60 | ~ spl30_64), inference(avatar_split_clause, [], [f7105, f637, f620])).
fof(f7105, plain, (~ (e3 = op(e0, e1)) | ~ spl30_64), inference(backward_demodulation, [], [f187, f639])).
fof(f7115, plain, (~ spl30_16 | ~ spl30_64), inference(avatar_split_clause, [], [f7104, f637, f433])).
fof(f7104, plain, (~ (e3 = op(e3, e0)) | ~ spl30_64), inference(backward_demodulation, [], [f165, f639])).
fof(f7075, plain, (~ spl30_37 | ~ spl30_47 | spl30_185), inference(avatar_split_clause, [], [f7074, f1258, f565, f523])).
fof(f7074, plain, (~ (e0 = op(e1, e2)) | (~ spl30_47 | spl30_185)), inference(forward_demodulation, [], [f1260, f567])).
fof(f7065, plain, (~ spl30_51 | spl30_207 | ~ spl30_210), inference(avatar_split_clause, [], [f7064, f1382, f1367, f582])).
fof(f1367, plain, (spl30_207 <=> (e2 = op(e0, op(e0, op(e0, e0))))), introduced(avatar_definition, [new_symbols(naming, [spl30_207])])).
fof(f7064, plain, (~ (e2 = op(e0, e3)) | (spl30_207 | ~ spl30_210)), inference(backward_demodulation, [], [f1369, f1383])).
fof(f1383, plain, ((e3 = op(e0, op(e0, e0))) | ~ spl30_210), inference(avatar_component_clause, [], [f1382])).
fof(f1369, plain, (~ (e2 = op(e0, op(e0, op(e0, e0)))) | spl30_207), inference(avatar_component_clause, [], [f1367])).
fof(f7045, plain, (~ spl30_3 | ~ spl30_6 | ~ spl30_9 | spl30_227), inference(avatar_contradiction_clause, [], [f7044])).
fof(f7044, plain, ($false | (~ spl30_3 | ~ spl30_6 | ~ spl30_9 | spl30_227)), inference(subsumption_resolution, [], [f7039, f406])).
fof(f7039, plain, (~ (e0 = op(e3, e1)) | (~ spl30_3 | ~ spl30_6 | spl30_227)), inference(backward_demodulation, [], [f6956, f393])).
fof(f6956, plain, (~ (e0 = op(e3, op(e3, e2))) | (~ spl30_3 | spl30_227)), inference(forward_demodulation, [], [f1469, f380])).
fof(f380, plain, ((e2 = op(e3, e3)) | ~ spl30_3), inference(avatar_component_clause, [], [f378])).
fof(f7042, plain, (~ spl30_6 | ~ spl30_42 | spl30_68), inference(avatar_contradiction_clause, [], [f7041])).
fof(f7041, plain, ($false | (~ spl30_6 | ~ spl30_42 | spl30_68)), inference(subsumption_resolution, [], [f7035, f546])).
fof(f7035, plain, (~ (e1 = op(e1, e1)) | (~ spl30_6 | spl30_68)), inference(backward_demodulation, [], [f690, f393])).
fof(f690, plain, (~ (e1 = op(e1, op(e3, e2))) | spl30_68), inference(avatar_component_clause, [], [f688])).
fof(f688, plain, (spl30_68 <=> (e1 = op(e1, op(e3, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl30_68])])).
fof(f7030, plain, (~ spl30_9 | ~ spl30_29 | ~ spl30_73), inference(avatar_contradiction_clause, [], [f7029])).
fof(f7029, plain, ($false | (~ spl30_9 | ~ spl30_29 | ~ spl30_73)), inference(subsumption_resolution, [], [f7028, f212])).
fof(f7028, plain, ((e0 = e2) | (~ spl30_9 | ~ spl30_29 | ~ spl30_73)), inference(forward_demodulation, [], [f7022, f491])).
fof(f7022, plain, ((e2 = op(e2, e0)) | (~ spl30_9 | ~ spl30_73)), inference(backward_demodulation, [], [f713, f406])).
fof(f7027, plain, (~ spl30_5 | ~ spl30_9), inference(avatar_split_clause, [], [f7017, f404, f387])).
fof(f7017, plain, (~ (e0 = op(e3, e2)) | ~ spl30_9), inference(backward_demodulation, [], [f208, f406])).
fof(f208, plain, ~ (op(e3, e1) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f7016, plain, (~ spl30_37 | spl30_54 | ~ spl30_177), inference(avatar_contradiction_clause, [], [f7015])).
fof(f7015, plain, ($false | (~ spl30_37 | spl30_54 | ~ spl30_177)), inference(subsumption_resolution, [], [f7009, f596])).
fof(f596, plain, (~ (e1 = op(e0, e2)) | spl30_54), inference(avatar_component_clause, [], [f595])).
fof(f7009, plain, ((e1 = op(e0, e2)) | (~ spl30_37 | ~ spl30_177)), inference(backward_demodulation, [], [f1219, f525])).
fof(f1219, plain, ((e1 = op(op(e1, e2), e2)) | ~ spl30_177), inference(avatar_component_clause, [], [f1218])).
fof(f7011, plain, (~ spl30_5 | ~ spl30_37), inference(avatar_split_clause, [], [f7001, f523, f387])).
fof(f7001, plain, (~ (e0 = op(e3, e2)) | ~ spl30_37), inference(backward_demodulation, [], [f179, f525])).
fof(f6994, plain, (~ spl30_10 | ~ spl30_42), inference(avatar_split_clause, [], [f6985, f544, f408])).
fof(f6985, plain, (~ (e1 = op(e3, e1)) | ~ spl30_42), inference(backward_demodulation, [], [f173, f546])).
fof(f6963, plain, (~ spl30_5 | ~ spl30_3 | spl30_223), inference(avatar_split_clause, [], [f6962, f1448, f378, f387])).
fof(f6962, plain, (~ (e0 = op(e3, e2)) | (~ spl30_3 | spl30_223)), inference(forward_demodulation, [], [f1450, f380])).
fof(f1450, plain, (~ (e0 = op(e3, op(e3, e3))) | spl30_223), inference(avatar_component_clause, [], [f1448])).
fof(f6947, plain, (~ spl30_18 | ~ spl30_24 | spl30_220), inference(avatar_contradiction_clause, [], [f6946])).
fof(f6946, plain, ($false | (~ spl30_18 | ~ spl30_24 | spl30_220)), inference(subsumption_resolution, [], [f6945, f444])).
fof(f6945, plain, (~ (e1 = op(e2, e3)) | (~ spl30_24 | spl30_220)), inference(forward_demodulation, [], [f1434, f469])).
fof(f1434, plain, (~ (e1 = op(e2, op(e2, e2))) | spl30_220), inference(avatar_component_clause, [], [f1432])).
fof(f6940, plain, (~ spl30_42 | ~ spl30_18 | spl30_86), inference(avatar_split_clause, [], [f6874, f775, f442, f544])).
fof(f775, plain, (spl30_86 <=> (e1 = op(e1, op(e2, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl30_86])])).
fof(f6874, plain, (~ (e1 = op(e1, e1)) | (~ spl30_18 | spl30_86)), inference(backward_demodulation, [], [f777, f444])).
fof(f777, plain, (~ (e1 = op(e1, op(e2, e3))) | spl30_86), inference(avatar_component_clause, [], [f775])).
fof(f6935, plain, (spl30_40 | ~ spl30_18 | ~ spl30_160), inference(avatar_split_clause, [], [f6877, f1136, f442, f535])).
fof(f6877, plain, ((e3 = op(e1, e2)) | (~ spl30_18 | ~ spl30_160)), inference(backward_demodulation, [], [f1138, f444])).
fof(f6910, plain, (~ spl30_2 | ~ spl30_10), inference(avatar_split_clause, [], [f6902, f408, f374])).
fof(f6902, plain, (~ (e1 = op(e3, e3)) | ~ spl30_10), inference(backward_demodulation, [], [f209, f410])).
fof(f410, plain, ((e1 = op(e3, e1)) | ~ spl30_10), inference(avatar_component_clause, [], [f408])).
fof(f6896, plain, (~ spl30_12 | ~ spl30_16), inference(avatar_split_clause, [], [f6888, f433, f416])).
fof(f6888, plain, (~ (e3 = op(e3, e1)) | ~ spl30_16), inference(backward_demodulation, [], [f205, f435])).
fof(f205, plain, ~ (op(e3, e0) = op(e3, e1)), inference(cnf_transformation, [], [f3])).
fof(f6882, plain, (~ spl30_18 | ~ spl30_27 | spl30_85), inference(avatar_contradiction_clause, [], [f6881])).
fof(f6881, plain, ($false | (~ spl30_18 | ~ spl30_27 | spl30_85)), inference(subsumption_resolution, [], [f6873, f482])).
fof(f6873, plain, (~ (e2 = op(e2, e1)) | (~ spl30_18 | spl30_85)), inference(backward_demodulation, [], [f772, f444])).
fof(f772, plain, (~ (e2 = op(e2, op(e2, e3))) | spl30_85), inference(avatar_component_clause, [], [f770])).
fof(f770, plain, (spl30_85 <=> (e2 = op(e2, op(e2, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl30_85])])).
fof(f6880, plain, (~ spl30_12 | ~ spl30_18 | spl30_84), inference(avatar_split_clause, [], [f6872, f765, f442, f416])).
fof(f765, plain, (spl30_84 <=> (e3 = op(e3, op(e2, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl30_84])])).
fof(f6872, plain, (~ (e3 = op(e3, e1)) | (~ spl30_18 | spl30_84)), inference(backward_demodulation, [], [f767, f444])).
fof(f767, plain, (~ (e3 = op(e3, op(e2, e3))) | spl30_84), inference(avatar_component_clause, [], [f765])).
fof(f6879, plain, (~ spl30_2 | ~ spl30_18), inference(avatar_split_clause, [], [f6868, f442, f374])).
fof(f6868, plain, (~ (e1 = op(e3, e3)) | ~ spl30_18), inference(backward_demodulation, [], [f186, f444])).
fof(f6846, plain, (~ spl30_25 | ~ spl30_29), inference(avatar_split_clause, [], [f6836, f489, f472])).
fof(f6836, plain, (~ (e0 = op(e2, e1)) | ~ spl30_29), inference(backward_demodulation, [], [f199, f491])).
fof(f6845, plain, (~ spl30_13 | ~ spl30_29), inference(avatar_split_clause, [], [f6835, f489, f421])).
fof(f6835, plain, (~ (e0 = op(e3, e0)) | ~ spl30_29), inference(backward_demodulation, [], [f168, f491])).
fof(f6831, plain, (~ spl30_12 | ~ spl30_36 | spl30_76), inference(avatar_split_clause, [], [f6824, f727, f518, f416])).
fof(f6824, plain, (~ (e3 = op(e3, e1)) | (~ spl30_36 | spl30_76)), inference(backward_demodulation, [], [f728, f520])).
fof(f6817, plain, (~ spl30_22 | ~ spl30_38), inference(avatar_split_clause, [], [f6805, f527, f459])).
fof(f6805, plain, (~ (e1 = op(e2, e2)) | ~ spl30_38), inference(backward_demodulation, [], [f178, f529])).
fof(f529, plain, ((e1 = op(e1, e2)) | ~ spl30_38), inference(avatar_component_clause, [], [f527])).
fof(f6785, plain, (~ spl30_49 | spl30_187), inference(avatar_contradiction_clause, [], [f6784])).
fof(f6784, plain, ($false | (~ spl30_49 | spl30_187)), inference(subsumption_resolution, [], [f6777, f576])).
fof(f576, plain, ((e0 = op(e0, e3)) | ~ spl30_49), inference(avatar_component_clause, [], [f574])).
fof(f6777, plain, (~ (e0 = op(e0, e3)) | (~ spl30_49 | spl30_187)), inference(backward_demodulation, [], [f1269, f576])).
fof(f6781, plain, (~ spl30_13 | ~ spl30_49 | spl30_82), inference(avatar_split_clause, [], [f6773, f756, f574, f421])).
fof(f6773, plain, (~ (e0 = op(e3, e0)) | (~ spl30_49 | spl30_82)), inference(backward_demodulation, [], [f757, f576])).
fof(f757, plain, (~ (op(e0, e3) = op(e3, e0)) | spl30_82), inference(avatar_component_clause, [], [f756])).
fof(f6780, plain, (~ spl30_33 | ~ spl30_49), inference(avatar_split_clause, [], [f6767, f574, f506])).
fof(f6767, plain, (~ (e0 = op(e1, e3)) | ~ spl30_49), inference(backward_demodulation, [], [f181, f576])).
fof(f6750, plain, (~ spl30_52 | ~ spl30_60), inference(avatar_split_clause, [], [f6743, f620, f586])).
fof(f6743, plain, (~ (e3 = op(e0, e3)) | ~ spl30_60), inference(backward_demodulation, [], [f191, f622])).
fof(f6732, plain, (~ spl30_42 | ~ spl30_62 | spl30_144), inference(avatar_split_clause, [], [f6731, f1060, f629, f544])).
fof(f629, plain, (spl30_62 <=> (op(e0, e0) = e1)), introduced(avatar_definition, [new_symbols(naming, [spl30_62])])).
fof(f6731, plain, (~ (e1 = op(e1, e1)) | (~ spl30_62 | spl30_144)), inference(forward_demodulation, [], [f1062, f631])).
fof(f631, plain, ((op(e0, e0) = e1) | ~ spl30_62), inference(avatar_component_clause, [], [f629])).
fof(f6730, plain, (~ spl30_25 | ~ spl30_47 | spl30_184), inference(avatar_split_clause, [], [f6729, f1253, f565, f472])).
fof(f6729, plain, (~ (e0 = op(e2, e1)) | (~ spl30_47 | spl30_184)), inference(forward_demodulation, [], [f1254, f567])).
fof(f6698, plain, (~ spl30_12 | ~ spl30_62 | spl30_142), inference(avatar_split_clause, [], [f6623, f1050, f629, f416])).
fof(f6623, plain, (~ (e3 = op(e3, e1)) | (~ spl30_62 | spl30_142)), inference(backward_demodulation, [], [f1052, f631])).
fof(f6674, plain, (~ spl30_3 | ~ spl30_19), inference(avatar_split_clause, [], [f6663, f446, f378])).
fof(f6663, plain, (~ (e2 = op(e3, e3)) | ~ spl30_19), inference(backward_demodulation, [], [f186, f448])).
fof(f448, plain, ((e2 = op(e2, e3)) | ~ spl30_19), inference(avatar_component_clause, [], [f446])).
fof(f6662, plain, (~ spl30_22 | ~ spl30_42 | spl30_91), inference(avatar_contradiction_clause, [], [f6661])).
fof(f6661, plain, ($false | (~ spl30_22 | ~ spl30_42 | spl30_91)), inference(subsumption_resolution, [], [f6655, f546])).
fof(f6655, plain, (~ (e1 = op(e1, e1)) | (~ spl30_22 | spl30_91)), inference(backward_demodulation, [], [f802, f461])).
fof(f802, plain, (~ (e1 = op(e1, op(e2, e2))) | spl30_91), inference(avatar_component_clause, [], [f800])).
fof(f800, plain, (spl30_91 <=> (e1 = op(e1, op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl30_91])])).
fof(f6660, plain, (~ spl30_12 | ~ spl30_22 | spl30_89), inference(avatar_contradiction_clause, [], [f6659])).
fof(f6659, plain, ($false | (~ spl30_12 | ~ spl30_22 | spl30_89)), inference(subsumption_resolution, [], [f6654, f418])).
fof(f418, plain, ((e3 = op(e3, e1)) | ~ spl30_12), inference(avatar_component_clause, [], [f416])).
fof(f6654, plain, (~ (e3 = op(e3, e1)) | (~ spl30_22 | spl30_89)), inference(backward_demodulation, [], [f792, f461])).
fof(f6658, plain, (~ spl30_18 | ~ spl30_22), inference(avatar_split_clause, [], [f6651, f459, f442])).
fof(f6651, plain, (~ (e1 = op(e2, e3)) | ~ spl30_22), inference(backward_demodulation, [], [f204, f461])).
fof(f6630, plain, (~ spl30_54 | ~ spl30_62), inference(avatar_split_clause, [], [f6620, f629, f595])).
fof(f6620, plain, (~ (e1 = op(e0, e2)) | ~ spl30_62), inference(backward_demodulation, [], [f188, f631])).
fof(f188, plain, ~ (op(e0, e0) = op(e0, e2)), inference(cnf_transformation, [], [f3])).
fof(f6629, plain, (~ spl30_14 | ~ spl30_62), inference(avatar_split_clause, [], [f6619, f629, f425])).
fof(f6619, plain, (~ (e1 = op(e3, e0)) | ~ spl30_62), inference(backward_demodulation, [], [f165, f631])).
fof(f6602, plain, (spl30_19 | ~ spl30_40 | ~ spl30_112), inference(avatar_split_clause, [], [f6601, f902, f535, f446])).
fof(f902, plain, (spl30_112 <=> (e2 = op(e2, op(e1, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl30_112])])).
fof(f6601, plain, ((e2 = op(e2, e3)) | (~ spl30_40 | ~ spl30_112)), inference(forward_demodulation, [], [f903, f537])).
fof(f903, plain, ((e2 = op(e2, op(e1, e2))) | ~ spl30_112), inference(avatar_component_clause, [], [f902])).
fof(f6582, plain, (~ spl30_54 | ~ spl30_25 | spl30_166), inference(avatar_split_clause, [], [f6581, f1166, f472, f595])).
fof(f6581, plain, (~ (e1 = op(e0, e2)) | (~ spl30_25 | spl30_166)), inference(forward_demodulation, [], [f1167, f474])).
fof(f474, plain, ((e0 = op(e2, e1)) | ~ spl30_25), inference(avatar_component_clause, [], [f472])).
fof(f1167, plain, (~ (e1 = op(op(e2, e1), e2)) | spl30_166), inference(avatar_component_clause, [], [f1166])).
fof(f6579, plain, (~ spl30_5 | ~ spl30_32 | spl30_170), inference(avatar_split_clause, [], [f6578, f1185, f501, f387])).
fof(f1185, plain, (spl30_170 <=> (e0 = op(op(e2, e0), e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_170])])).
fof(f6578, plain, (~ (e0 = op(e3, e2)) | (~ spl30_32 | spl30_170)), inference(forward_demodulation, [], [f1186, f503])).
fof(f1186, plain, (~ (e0 = op(op(e2, e0), e2)) | spl30_170), inference(avatar_component_clause, [], [f1185])).
fof(f6577, plain, (~ spl30_54 | ~ spl30_59 | spl30_197), inference(avatar_split_clause, [], [f6576, f1315, f616, f595])).
fof(f6576, plain, (~ (e1 = op(e0, e2)) | (~ spl30_59 | spl30_197)), inference(forward_demodulation, [], [f1317, f618])).
fof(f6572, plain, (~ spl30_54 | ~ spl30_208 | spl30_209), inference(avatar_split_clause, [], [f6569, f1377, f1372, f595])).
fof(f6569, plain, (~ (e1 = op(e0, e2)) | (~ spl30_208 | spl30_209)), inference(backward_demodulation, [], [f1379, f1373])).
fof(f1373, plain, ((e2 = op(e0, op(e0, e0))) | ~ spl30_208), inference(avatar_component_clause, [], [f1372])).
fof(f6548, plain, (~ spl30_63 | ~ spl30_47), inference(avatar_split_clause, [], [f6547, f565, f633])).
fof(f6547, plain, (~ (op(e0, e0) = e2) | ~ spl30_47), inference(forward_demodulation, [], [f163, f567])).
fof(f163, plain, ~ (op(e0, e0) = op(e1, e0)), inference(cnf_transformation, [], [f3])).
fof(f6541, plain, (~ spl30_16 | ~ spl30_32), inference(avatar_split_clause, [], [f6540, f501, f433])).
fof(f6540, plain, (~ (e3 = op(e3, e0)) | ~ spl30_32), inference(forward_demodulation, [], [f168, f503])).
fof(f6539, plain, (~ spl30_15 | ~ spl30_47), inference(avatar_split_clause, [], [f6538, f565, f429])).
fof(f6538, plain, (~ (e2 = op(e3, e0)) | ~ spl30_47), inference(forward_demodulation, [], [f167, f567])).
fof(f6535, plain, (~ spl30_15 | ~ spl30_32 | spl30_169), inference(avatar_split_clause, [], [f6223, f1180, f501, f429])).
fof(f6223, plain, (~ (e2 = op(e3, e0)) | (~ spl30_32 | spl30_169)), inference(forward_demodulation, [], [f1182, f503])).
fof(f6530, plain, (~ spl30_16 | ~ spl30_25 | spl30_94), inference(avatar_split_clause, [], [f6504, f814, f472, f433])).
fof(f6504, plain, (~ (e3 = op(e3, e0)) | (~ spl30_25 | spl30_94)), inference(backward_demodulation, [], [f816, f474])).
fof(f6529, plain, (~ spl30_8 | ~ spl30_40), inference(avatar_split_clause, [], [f6528, f535, f399])).
fof(f6528, plain, (~ (e3 = op(e3, e2)) | ~ spl30_40), inference(forward_demodulation, [], [f179, f537])).
fof(f6527, plain, (~ spl30_8 | ~ spl30_12), inference(avatar_split_clause, [], [f6526, f416, f399])).
fof(f6526, plain, (~ (e3 = op(e3, e2)) | ~ spl30_12), inference(forward_demodulation, [], [f208, f418])).
fof(f6524, plain, (~ spl30_4 | ~ spl30_52), inference(avatar_split_clause, [], [f6523, f586, f382])).
fof(f6523, plain, (~ (e3 = op(e3, e3)) | ~ spl30_52), inference(forward_demodulation, [], [f183, f588])).
fof(f588, plain, ((e3 = op(e0, e3)) | ~ spl30_52), inference(avatar_component_clause, [], [f586])).
fof(f6511, plain, (~ spl30_25 | spl30_30 | ~ spl30_167), inference(avatar_contradiction_clause, [], [f6510])).
fof(f6510, plain, ($false | (~ spl30_25 | spl30_30 | ~ spl30_167)), inference(subsumption_resolution, [], [f6507, f494])).
fof(f494, plain, (~ (e1 = op(e2, e0)) | spl30_30), inference(avatar_component_clause, [], [f493])).
fof(f6507, plain, ((e1 = op(e2, e0)) | (~ spl30_25 | ~ spl30_167)), inference(backward_demodulation, [], [f1172, f474])).
fof(f6508, plain, (~ spl30_17 | ~ spl30_25), inference(avatar_split_clause, [], [f6500, f472, f438])).
fof(f6500, plain, (~ (e0 = op(e2, e3)) | ~ spl30_25), inference(backward_demodulation, [], [f203, f474])).
fof(f6497, plain, (~ spl30_33 | ~ spl30_61 | spl30_109), inference(avatar_contradiction_clause, [], [f6496])).
fof(f6496, plain, ($false | (~ spl30_33 | ~ spl30_61 | spl30_109)), inference(subsumption_resolution, [], [f6490, f627])).
fof(f6490, plain, (~ (e0 = op(e0, e0)) | (~ spl30_33 | spl30_109)), inference(backward_demodulation, [], [f889, f508])).
fof(f889, plain, (~ (e0 = op(e0, op(e1, e3))) | spl30_109), inference(avatar_component_clause, [], [f887])).
fof(f6468, plain, (~ spl30_42 | ~ spl30_214), inference(avatar_contradiction_clause, [], [f6467])).
fof(f6467, plain, ($false | (~ spl30_42 | ~ spl30_214)), inference(subsumption_resolution, [], [f6466, f214])).
fof(f214, plain, ~ (e1 = e2), inference(cnf_transformation, [], [f4])).
fof(f6466, plain, ((e1 = e2) | (~ spl30_42 | ~ spl30_214)), inference(forward_demodulation, [], [f6458, f546])).
fof(f6458, plain, ((e2 = op(e1, e1)) | (~ spl30_42 | ~ spl30_214)), inference(backward_demodulation, [], [f1403, f546])).
fof(f6464, plain, (~ spl30_42 | spl30_118), inference(avatar_contradiction_clause, [], [f6463])).
fof(f6463, plain, ($false | (~ spl30_42 | spl30_118)), inference(subsumption_resolution, [], [f6455, f546])).
fof(f6455, plain, (~ (e1 = op(e1, e1)) | (~ spl30_42 | spl30_118)), inference(backward_demodulation, [], [f934, f546])).
fof(f934, plain, (~ (e1 = op(e1, op(e1, e1))) | spl30_118), inference(avatar_component_clause, [], [f932])).
fof(f932, plain, (spl30_118 <=> (e1 = op(e1, op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl30_118])])).
fof(f6443, plain, (~ spl30_36 | ~ spl30_52), inference(avatar_split_clause, [], [f6434, f586, f518])).
fof(f6434, plain, (~ (e3 = op(e1, e3)) | ~ spl30_52), inference(backward_demodulation, [], [f181, f588])).
fof(f6431, plain, (~ spl30_38 | ~ spl30_54), inference(avatar_split_clause, [], [f6420, f595, f527])).
fof(f6420, plain, (~ (e1 = op(e1, e2)) | ~ spl30_54), inference(backward_demodulation, [], [f175, f597])).
fof(f6407, plain, (spl30_25 | ~ spl30_47 | ~ spl30_184), inference(avatar_split_clause, [], [f6406, f1253, f565, f472])).
fof(f6406, plain, ((e0 = op(e2, e1)) | (~ spl30_47 | ~ spl30_184)), inference(forward_demodulation, [], [f1255, f567])).
fof(f1255, plain, ((e0 = op(op(e1, e0), e1)) | ~ spl30_184), inference(avatar_component_clause, [], [f1253])).
fof(f6403, plain, (spl30_54 | ~ spl30_59 | ~ spl30_197), inference(avatar_split_clause, [], [f6402, f1315, f616, f595])).
fof(f6402, plain, ((e1 = op(e0, e2)) | (~ spl30_59 | ~ spl30_197)), inference(forward_demodulation, [], [f1316, f618])).
fof(f1316, plain, ((e1 = op(e0, op(e0, e1))) | ~ spl30_197), inference(avatar_component_clause, [], [f1315])).
fof(f6391, plain, (~ spl30_3 | ~ spl30_5 | ~ spl30_14 | spl30_225), inference(avatar_contradiction_clause, [], [f6390])).
fof(f6390, plain, ($false | (~ spl30_3 | ~ spl30_5 | ~ spl30_14 | spl30_225)), inference(subsumption_resolution, [], [f6389, f427])).
fof(f6389, plain, (~ (e1 = op(e3, e0)) | (~ spl30_3 | ~ spl30_5 | spl30_225)), inference(forward_demodulation, [], [f6388, f389])).
fof(f389, plain, ((e0 = op(e3, e2)) | ~ spl30_5), inference(avatar_component_clause, [], [f387])).
fof(f6388, plain, (~ (e1 = op(e3, op(e3, e2))) | (~ spl30_3 | spl30_225)), inference(forward_demodulation, [], [f1459, f380])).
fof(f6370, plain, (~ spl30_51 | ~ spl30_59), inference(avatar_split_clause, [], [f6369, f616, f582])).
fof(f6369, plain, (~ (e2 = op(e0, e3)) | ~ spl30_59), inference(forward_demodulation, [], [f191, f618])).
fof(f6363, plain, (~ spl30_43 | ~ spl30_47), inference(avatar_split_clause, [], [f6362, f565, f548])).
fof(f6362, plain, (~ (e2 = op(e1, e1)) | ~ spl30_47), inference(forward_demodulation, [], [f193, f567])).
fof(f6314, plain, (~ spl30_19 | ~ spl30_23), inference(avatar_split_clause, [], [f6308, f463, f446])).
fof(f6308, plain, (~ (e2 = op(e2, e3)) | ~ spl30_23), inference(backward_demodulation, [], [f204, f465])).
fof(f6313, plain, (~ spl30_7 | ~ spl30_23), inference(avatar_split_clause, [], [f6307, f463, f395])).
fof(f6307, plain, (~ (e2 = op(e3, e2)) | ~ spl30_23), inference(backward_demodulation, [], [f180, f465])).
fof(f6286, plain, (spl30_5 | ~ spl30_32 | ~ spl30_170), inference(avatar_split_clause, [], [f6285, f1185, f501, f387])).
fof(f6285, plain, ((e0 = op(e3, e2)) | (~ spl30_32 | ~ spl30_170)), inference(forward_demodulation, [], [f1187, f503])).
fof(f1187, plain, ((e0 = op(op(e2, e0), e2)) | ~ spl30_170), inference(avatar_component_clause, [], [f1185])).
fof(f6268, plain, (~ spl30_4 | ~ spl30_12), inference(avatar_split_clause, [], [f6267, f416, f382])).
fof(f6267, plain, (~ (e3 = op(e3, e3)) | ~ spl30_12), inference(forward_demodulation, [], [f209, f418])).
fof(f6248, plain, (~ spl30_21 | ~ spl30_61 | spl30_92), inference(avatar_contradiction_clause, [], [f6247])).
fof(f6247, plain, ($false | (~ spl30_21 | ~ spl30_61 | spl30_92)), inference(subsumption_resolution, [], [f6246, f627])).
fof(f6246, plain, (~ (e0 = op(e0, e0)) | (~ spl30_21 | spl30_92)), inference(forward_demodulation, [], [f807, f457])).
fof(f807, plain, (~ (e0 = op(e0, op(e2, e2))) | spl30_92), inference(avatar_component_clause, [], [f805])).
fof(f6233, plain, (~ spl30_21 | ~ spl30_59 | ~ spl30_138), inference(avatar_contradiction_clause, [], [f6232])).
fof(f6232, plain, ($false | (~ spl30_21 | ~ spl30_59 | ~ spl30_138)), inference(subsumption_resolution, [], [f6231, f212])).
fof(f6231, plain, ((e0 = e2) | (~ spl30_21 | ~ spl30_59 | ~ spl30_138)), inference(forward_demodulation, [], [f6230, f457])).
fof(f6230, plain, ((e2 = op(e2, e2)) | (~ spl30_59 | ~ spl30_138)), inference(forward_demodulation, [], [f1031, f618])).
fof(f1031, plain, ((e2 = op(e2, op(e0, e1))) | ~ spl30_138), inference(avatar_component_clause, [], [f1030])).
fof(f6227, plain, (~ spl30_14 | spl30_33 | ~ spl30_156), inference(avatar_contradiction_clause, [], [f6226])).
fof(f6226, plain, ($false | (~ spl30_14 | spl30_33 | ~ spl30_156)), inference(subsumption_resolution, [], [f6225, f507])).
fof(f507, plain, (~ (e0 = op(e1, e3)) | spl30_33), inference(avatar_component_clause, [], [f506])).
fof(f6225, plain, ((e0 = op(e1, e3)) | (~ spl30_14 | ~ spl30_156)), inference(forward_demodulation, [], [f1119, f427])).
fof(f1119, plain, ((e0 = op(op(e3, e0), e3)) | ~ spl30_156), inference(avatar_component_clause, [], [f1117])).
fof(f1117, plain, (spl30_156 <=> (e0 = op(op(e3, e0), e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_156])])).
fof(f6202, plain, (~ spl30_11 | ~ spl30_19 | ~ spl30_70), inference(avatar_split_clause, [], [f6196, f698, f446, f412])).
fof(f6196, plain, (~ (e2 = op(e3, e1)) | (~ spl30_19 | ~ spl30_70)), inference(backward_demodulation, [], [f6068, f448])).
fof(f6068, plain, (~ (op(e2, e3) = op(e3, e1)) | ~ spl30_70), inference(backward_demodulation, [], [f208, f700])).
fof(f6187, plain, (spl30_17 | ~ spl30_32 | ~ spl30_171), inference(avatar_contradiction_clause, [], [f6186])).
fof(f6186, plain, ($false | (spl30_17 | ~ spl30_32 | ~ spl30_171)), inference(subsumption_resolution, [], [f6181, f439])).
fof(f6181, plain, ((e0 = op(e2, e3)) | (~ spl30_32 | ~ spl30_171)), inference(backward_demodulation, [], [f1191, f503])).
fof(f1191, plain, ((e0 = op(e2, op(e2, e0))) | ~ spl30_171), inference(avatar_component_clause, [], [f1190])).
fof(f6167, plain, (~ spl30_35 | ~ spl30_47), inference(avatar_split_clause, [], [f6160, f565, f514])).
fof(f6160, plain, (~ (e2 = op(e1, e3)) | ~ spl30_47), inference(backward_demodulation, [], [f195, f567])).
fof(f6156, plain, (spl30_34 | ~ spl30_56 | ~ spl30_134), inference(avatar_contradiction_clause, [], [f6155])).
fof(f6155, plain, ($false | (spl30_34 | ~ spl30_56 | ~ spl30_134)), inference(subsumption_resolution, [], [f6147, f511])).
fof(f511, plain, (~ (e1 = op(e1, e3)) | spl30_34), inference(avatar_component_clause, [], [f510])).
fof(f6147, plain, ((e1 = op(e1, e3)) | (~ spl30_56 | ~ spl30_134)), inference(backward_demodulation, [], [f1011, f605])).
fof(f6140, plain, (~ spl30_11 | ~ spl30_59), inference(avatar_split_clause, [], [f6131, f616, f412])).
fof(f6131, plain, (~ (e2 = op(e3, e1)) | ~ spl30_59), inference(backward_demodulation, [], [f171, f618])).
fof(f6125, plain, (~ spl30_31 | ~ spl30_21 | spl30_90), inference(avatar_split_clause, [], [f6124, f795, f455, f497])).
fof(f6124, plain, (~ (e2 = op(e2, e0)) | (~ spl30_21 | spl30_90)), inference(forward_demodulation, [], [f797, f457])).
fof(f797, plain, (~ (e2 = op(e2, op(e2, e2))) | spl30_90), inference(avatar_component_clause, [], [f795])).
fof(f6123, plain, (spl30_12 | ~ spl30_26 | ~ spl30_94), inference(avatar_split_clause, [], [f6122, f814, f476, f416])).
fof(f6122, plain, ((e3 = op(e3, e1)) | (~ spl30_26 | ~ spl30_94)), inference(forward_demodulation, [], [f815, f478])).
fof(f815, plain, ((e3 = op(e3, op(e2, e1))) | ~ spl30_94), inference(avatar_component_clause, [], [f814])).
fof(f6109, plain, (~ spl30_60 | ~ spl30_50 | spl30_189), inference(avatar_split_clause, [], [f6108, f1277, f578, f620])).
fof(f6108, plain, (~ (e3 = op(e0, e1)) | (~ spl30_50 | spl30_189)), inference(forward_demodulation, [], [f1279, f580])).
fof(f6107, plain, (~ spl30_31 | ~ spl30_1 | spl30_201), inference(avatar_split_clause, [], [f6106, f1336, f370, f497])).
fof(f6106, plain, (~ (e2 = op(e2, e0)) | (~ spl30_1 | spl30_201)), inference(forward_demodulation, [], [f1338, f372])).
fof(f372, plain, ((e0 = op(e3, e3)) | ~ spl30_1), inference(avatar_component_clause, [], [f370])).
fof(f1338, plain, (~ (e2 = op(e2, op(e3, e3))) | spl30_201), inference(avatar_component_clause, [], [f1336])).
fof(f6065, plain, (~ spl30_12 | ~ spl30_50 | spl30_127), inference(avatar_split_clause, [], [f5949, f975, f578, f416])).
fof(f5949, plain, (~ (e3 = op(e3, e1)) | (~ spl30_50 | spl30_127)), inference(backward_demodulation, [], [f977, f580])).
fof(f6064, plain, (~ spl30_11 | ~ spl30_1 | ~ spl30_14 | spl30_224), inference(avatar_split_clause, [], [f6056, f1452, f425, f370, f412])).
fof(f1452, plain, (spl30_224 <=> (e2 = op(e3, op(e3, op(e3, e3))))), introduced(avatar_definition, [new_symbols(naming, [spl30_224])])).
fof(f6056, plain, (~ (e2 = op(e3, e1)) | (~ spl30_1 | ~ spl30_14 | spl30_224)), inference(forward_demodulation, [], [f6050, f427])).
fof(f6050, plain, (~ (e2 = op(e3, op(e3, e0))) | (~ spl30_1 | spl30_224)), inference(backward_demodulation, [], [f1454, f372])).
fof(f1454, plain, (~ (e2 = op(e3, op(e3, op(e3, e3)))) | spl30_224), inference(avatar_component_clause, [], [f1452])).
fof(f6061, plain, (~ spl30_1 | ~ spl30_14 | spl30_226), inference(avatar_contradiction_clause, [], [f6060])).
fof(f6060, plain, ($false | (~ spl30_1 | ~ spl30_14 | spl30_226)), inference(subsumption_resolution, [], [f6052, f427])).
fof(f6052, plain, (~ (e1 = op(e3, e0)) | (~ spl30_1 | spl30_226)), inference(backward_demodulation, [], [f1464, f372])).
fof(f1464, plain, (~ (e1 = op(e3, op(e3, e3))) | spl30_226), inference(avatar_component_clause, [], [f1462])).
fof(f6008, plain, (~ spl30_26 | spl30_167), inference(avatar_contradiction_clause, [], [f6007])).
fof(f6007, plain, ($false | (~ spl30_26 | spl30_167)), inference(subsumption_resolution, [], [f6004, f478])).
fof(f6004, plain, (~ (e1 = op(e2, e1)) | (~ spl30_26 | spl30_167)), inference(backward_demodulation, [], [f1173, f478])).
fof(f1173, plain, (~ (e1 = op(e2, op(e2, e1))) | spl30_167), inference(avatar_component_clause, [], [f1171])).
fof(f5992, plain, (~ spl30_35 | ~ spl30_38 | spl30_108), inference(avatar_contradiction_clause, [], [f5991])).
fof(f5991, plain, ($false | (~ spl30_35 | ~ spl30_38 | spl30_108)), inference(subsumption_resolution, [], [f5982, f529])).
fof(f5982, plain, (~ (e1 = op(e1, e2)) | (~ spl30_35 | spl30_108)), inference(backward_demodulation, [], [f884, f516])).
fof(f884, plain, (~ (e1 = op(e1, op(e1, e3))) | spl30_108), inference(avatar_component_clause, [], [f882])).
fof(f882, plain, (spl30_108 <=> (e1 = op(e1, op(e1, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl30_108])])).
fof(f5972, plain, (~ spl30_34 | ~ spl30_38), inference(avatar_split_clause, [], [f5961, f527, f510])).
fof(f5961, plain, (~ (e1 = op(e1, e3)) | ~ spl30_38), inference(backward_demodulation, [], [f198, f529])).
fof(f5954, plain, (~ spl30_34 | ~ spl30_50), inference(avatar_split_clause, [], [f5944, f578, f510])).
fof(f5944, plain, (~ (e1 = op(e1, e3)) | ~ spl30_50), inference(backward_demodulation, [], [f181, f580])).
fof(f5943, plain, (spl30_21 | ~ spl30_55 | ~ spl30_104 | ~ spl30_171), inference(avatar_split_clause, [], [f5936, f1190, f863, f599, f455])).
fof(f5936, plain, ((e0 = op(e2, e2)) | (~ spl30_55 | ~ spl30_104 | ~ spl30_171)), inference(backward_demodulation, [], [f5873, f601])).
fof(f5873, plain, ((e0 = op(e2, op(e0, e2))) | (~ spl30_104 | ~ spl30_171)), inference(backward_demodulation, [], [f1191, f865])).
fof(f5941, plain, (spl30_31 | ~ spl30_55 | ~ spl30_192), inference(avatar_split_clause, [], [f5932, f1291, f599, f497])).
fof(f5932, plain, ((e2 = op(e2, e0)) | (~ spl30_55 | ~ spl30_192)), inference(backward_demodulation, [], [f1293, f601])).
fof(f5940, plain, (spl30_21 | ~ spl30_55 | ~ spl30_191), inference(avatar_split_clause, [], [f5931, f1286, f599, f455])).
fof(f5931, plain, ((e0 = op(e2, e2)) | (~ spl30_55 | ~ spl30_191)), inference(backward_demodulation, [], [f1287, f601])).
fof(f1287, plain, ((e0 = op(op(e0, e2), e2)) | ~ spl30_191), inference(avatar_component_clause, [], [f1286])).
fof(f5938, plain, (spl30_31 | ~ spl30_55 | ~ spl30_104), inference(avatar_split_clause, [], [f5928, f863, f599, f497])).
fof(f5928, plain, ((e2 = op(e2, e0)) | (~ spl30_55 | ~ spl30_104)), inference(backward_demodulation, [], [f865, f601])).
fof(f5937, plain, (~ spl30_39 | ~ spl30_55), inference(avatar_split_clause, [], [f5923, f599, f531])).
fof(f5923, plain, (~ (e2 = op(e1, e2)) | ~ spl30_55), inference(backward_demodulation, [], [f175, f601])).
fof(f5922, plain, (spl30_50 | ~ spl30_60 | ~ spl30_197), inference(avatar_split_clause, [], [f5916, f1315, f620, f578])).
fof(f5916, plain, ((e1 = op(e0, e3)) | (~ spl30_60 | ~ spl30_197)), inference(backward_demodulation, [], [f1316, f622])).
fof(f5920, plain, (~ spl30_19 | ~ spl30_60 | spl30_138), inference(avatar_split_clause, [], [f5914, f1030, f620, f446])).
fof(f5914, plain, (~ (e2 = op(e2, e3)) | (~ spl30_60 | spl30_138)), inference(backward_demodulation, [], [f1032, f622])).
fof(f5917, plain, (~ spl30_28 | ~ spl30_60), inference(avatar_split_clause, [], [f5910, f620, f484])).
fof(f5910, plain, (~ (e3 = op(e2, e1)) | ~ spl30_60), inference(backward_demodulation, [], [f170, f622])).
fof(f5908, plain, (~ spl30_49 | ~ spl30_61), inference(avatar_split_clause, [], [f5901, f625, f574])).
fof(f5901, plain, (~ (e0 = op(e0, e3)) | ~ spl30_61), inference(backward_demodulation, [], [f189, f627])).
fof(f189, plain, ~ (op(e0, e0) = op(e0, e3)), inference(cnf_transformation, [], [f3])).
fof(f5907, plain, (~ spl30_29 | ~ spl30_61), inference(avatar_split_clause, [], [f5898, f625, f489])).
fof(f5898, plain, (~ (e0 = op(e2, e0)) | ~ spl30_61), inference(backward_demodulation, [], [f164, f627])).
fof(f5880, plain, (spl30_35 | ~ spl30_11 | ~ spl30_76), inference(avatar_split_clause, [], [f5879, f727, f412, f514])).
fof(f5879, plain, ((e2 = op(e1, e3)) | (~ spl30_11 | ~ spl30_76)), inference(forward_demodulation, [], [f729, f414])).
fof(f5875, plain, (~ spl30_104 | spl30_170 | ~ spl30_191), inference(avatar_contradiction_clause, [], [f5874])).
fof(f5874, plain, ($false | (~ spl30_104 | spl30_170 | ~ spl30_191)), inference(subsumption_resolution, [], [f5872, f1287])).
fof(f5872, plain, (~ (e0 = op(op(e0, e2), e2)) | (~ spl30_104 | spl30_170)), inference(backward_demodulation, [], [f1186, f865])).
fof(f5865, plain, (~ spl30_145 | ~ spl30_210), inference(avatar_contradiction_clause, [], [f5864])).
fof(f5864, plain, ($false | (~ spl30_145 | ~ spl30_210)), inference(subsumption_resolution, [], [f5863, f213])).
fof(f5863, plain, ((e0 = e3) | (~ spl30_145 | ~ spl30_210)), inference(backward_demodulation, [], [f1383, f1066])).
fof(f1066, plain, ((e0 = op(e0, op(e0, e0))) | ~ spl30_145), inference(avatar_component_clause, [], [f1065])).
fof(f1065, plain, (spl30_145 <=> (e0 = op(e0, op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl30_145])])).
fof(f5860, plain, (~ spl30_11 | spl30_18 | ~ spl30_152), inference(avatar_contradiction_clause, [], [f5859])).
fof(f5859, plain, ($false | (~ spl30_11 | spl30_18 | ~ spl30_152)), inference(subsumption_resolution, [], [f5858, f443])).
fof(f443, plain, (~ (e1 = op(e2, e3)) | spl30_18), inference(avatar_component_clause, [], [f442])).
fof(f5858, plain, ((e1 = op(e2, e3)) | (~ spl30_11 | ~ spl30_152)), inference(forward_demodulation, [], [f1100, f414])).
fof(f1100, plain, ((e1 = op(op(e3, e1), e3)) | ~ spl30_152), inference(avatar_component_clause, [], [f1098])).
fof(f5850, plain, (~ spl30_41 | ~ spl30_48 | spl30_216), inference(avatar_contradiction_clause, [], [f5849])).
fof(f5849, plain, ($false | (~ spl30_41 | ~ spl30_48 | spl30_216)), inference(subsumption_resolution, [], [f5848, f571])).
fof(f5848, plain, (~ (e3 = op(e1, e0)) | (~ spl30_41 | spl30_216)), inference(forward_demodulation, [], [f1414, f542])).
fof(f1414, plain, (~ (e3 = op(e1, op(e1, e1))) | spl30_216), inference(avatar_component_clause, [], [f1412])).
fof(f5836, plain, (~ spl30_50 | spl30_209 | ~ spl30_210), inference(avatar_split_clause, [], [f5492, f1382, f1377, f578])).
fof(f5492, plain, (~ (e1 = op(e0, e3)) | (spl30_209 | ~ spl30_210)), inference(forward_demodulation, [], [f1379, f1383])).
fof(f5834, plain, (~ spl30_38 | ~ spl30_11 | spl30_74), inference(avatar_split_clause, [], [f5798, f717, f412, f527])).
fof(f5798, plain, (~ (e1 = op(e1, e2)) | (~ spl30_11 | spl30_74)), inference(backward_demodulation, [], [f719, f414])).
fof(f5833, plain, (~ spl30_36 | ~ spl30_48), inference(avatar_split_clause, [], [f5832, f569, f518])).
fof(f5832, plain, (~ (e3 = op(e1, e3)) | ~ spl30_48), inference(forward_demodulation, [], [f195, f571])).
fof(f5823, plain, (spl30_28 | ~ spl30_11 | ~ spl30_151), inference(avatar_split_clause, [], [f5800, f1093, f412, f484])).
fof(f1093, plain, (spl30_151 <=> (e3 = op(op(e3, e1), e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_151])])).
fof(f5800, plain, ((e3 = op(e2, e1)) | (~ spl30_11 | ~ spl30_151)), inference(backward_demodulation, [], [f1094, f414])).
fof(f1094, plain, ((e3 = op(op(e3, e1), e1)) | ~ spl30_151), inference(avatar_component_clause, [], [f1093])).
fof(f5803, plain, (~ spl30_7 | ~ spl30_11), inference(avatar_split_clause, [], [f5795, f412, f395])).
fof(f5795, plain, (~ (e2 = op(e3, e2)) | ~ spl30_11), inference(backward_demodulation, [], [f208, f414])).
fof(f5783, plain, (~ spl30_41 | spl30_45 | ~ spl30_211), inference(avatar_contradiction_clause, [], [f5782])).
fof(f5782, plain, ($false | (~ spl30_41 | spl30_45 | ~ spl30_211)), inference(subsumption_resolution, [], [f5777, f558])).
fof(f558, plain, (~ (e0 = op(e1, e0)) | spl30_45), inference(avatar_component_clause, [], [f557])).
fof(f557, plain, (spl30_45 <=> (e0 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_45])])).
fof(f5777, plain, ((e0 = op(e1, e0)) | (~ spl30_41 | ~ spl30_211)), inference(backward_demodulation, [], [f1389, f542])).
fof(f1389, plain, ((e0 = op(e1, op(e1, e1))) | ~ spl30_211), inference(avatar_component_clause, [], [f1388])).
fof(f1388, plain, (spl30_211 <=> (e0 = op(e1, op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl30_211])])).
fof(f5754, plain, (~ spl30_58 | spl30_197), inference(avatar_contradiction_clause, [], [f5753])).
fof(f5753, plain, ($false | (~ spl30_58 | spl30_197)), inference(subsumption_resolution, [], [f5752, f614])).
fof(f5752, plain, (~ (e1 = op(e0, e1)) | (~ spl30_58 | spl30_197)), inference(forward_demodulation, [], [f1317, f614])).
fof(f5729, plain, (~ spl30_10 | ~ spl30_14), inference(avatar_split_clause, [], [f5728, f425, f408])).
fof(f5728, plain, (~ (e1 = op(e3, e1)) | ~ spl30_14), inference(forward_demodulation, [], [f205, f427])).
fof(f5727, plain, (spl30_9 | ~ spl30_48 | ~ spl30_184), inference(avatar_split_clause, [], [f5703, f1253, f569, f404])).
fof(f5703, plain, ((e0 = op(e3, e1)) | (~ spl30_48 | ~ spl30_184)), inference(forward_demodulation, [], [f1255, f571])).
fof(f5697, plain, (~ spl30_34 | ~ spl30_48 | spl30_123), inference(avatar_contradiction_clause, [], [f5696])).
fof(f5696, plain, ($false | (~ spl30_34 | ~ spl30_48 | spl30_123)), inference(subsumption_resolution, [], [f5695, f512])).
fof(f5695, plain, (~ (e1 = op(e1, e3)) | (~ spl30_48 | spl30_123)), inference(forward_demodulation, [], [f958, f571])).
fof(f958, plain, (~ (e1 = op(e1, op(e1, e0))) | spl30_123), inference(avatar_component_clause, [], [f956])).
fof(f956, plain, (spl30_123 <=> (e1 = op(e1, op(e1, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl30_123])])).
fof(f5693, plain, (~ spl30_14 | ~ spl30_48 | spl30_183), inference(avatar_contradiction_clause, [], [f5692])).
fof(f5692, plain, ($false | (~ spl30_14 | ~ spl30_48 | spl30_183)), inference(subsumption_resolution, [], [f5691, f427])).
fof(f5691, plain, (~ (e1 = op(e3, e0)) | (~ spl30_48 | spl30_183)), inference(forward_demodulation, [], [f1250, f571])).
fof(f1250, plain, (~ (e1 = op(op(e1, e0), e0)) | spl30_183), inference(avatar_component_clause, [], [f1248])).
fof(f1248, plain, (spl30_183 <=> (e1 = op(op(e1, e0), e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_183])])).
fof(f5690, plain, (~ spl30_37 | ~ spl30_43 | ~ spl30_48 | spl30_212), inference(avatar_contradiction_clause, [], [f5689])).
fof(f5689, plain, ($false | (~ spl30_37 | ~ spl30_43 | ~ spl30_48 | spl30_212)), inference(subsumption_resolution, [], [f5688, f571])).
fof(f5688, plain, (~ (e3 = op(e1, e0)) | (~ spl30_37 | ~ spl30_43 | spl30_212)), inference(forward_demodulation, [], [f5330, f525])).
fof(f5330, plain, (~ (e3 = op(e1, op(e1, e2))) | (~ spl30_43 | spl30_212)), inference(backward_demodulation, [], [f1394, f550])).
fof(f550, plain, ((e2 = op(e1, e1)) | ~ spl30_43), inference(avatar_component_clause, [], [f548])).
fof(f5621, plain, (~ spl30_29 | spl30_53 | ~ spl30_170), inference(avatar_contradiction_clause, [], [f5620])).
fof(f5620, plain, ($false | (~ spl30_29 | spl30_53 | ~ spl30_170)), inference(subsumption_resolution, [], [f5617, f592])).
fof(f5617, plain, ((e0 = op(e0, e2)) | (~ spl30_29 | ~ spl30_170)), inference(backward_demodulation, [], [f1187, f491])).
fof(f5619, plain, (~ spl30_17 | ~ spl30_29), inference(avatar_split_clause, [], [f5612, f489, f438])).
fof(f5612, plain, (~ (e0 = op(e2, e3)) | ~ spl30_29), inference(backward_demodulation, [], [f201, f491])).
fof(f5611, plain, (~ spl30_37 | ~ spl30_58 | ~ spl30_178), inference(avatar_contradiction_clause, [], [f5610])).
fof(f5610, plain, ($false | (~ spl30_37 | ~ spl30_58 | ~ spl30_178)), inference(subsumption_resolution, [], [f5609, f214])).
fof(f5609, plain, ((e1 = e2) | (~ spl30_37 | ~ spl30_58 | ~ spl30_178)), inference(forward_demodulation, [], [f5602, f614])).
fof(f5602, plain, ((e2 = op(e0, e1)) | (~ spl30_37 | ~ spl30_178)), inference(backward_demodulation, [], [f1225, f525])).
fof(f5606, plain, (~ spl30_31 | ~ spl30_37 | spl30_112), inference(avatar_split_clause, [], [f5598, f902, f523, f497])).
fof(f5598, plain, (~ (e2 = op(e2, e0)) | (~ spl30_37 | spl30_112)), inference(backward_demodulation, [], [f904, f525])).
fof(f904, plain, (~ (e2 = op(e2, op(e1, e2))) | spl30_112), inference(avatar_component_clause, [], [f902])).
fof(f5585, plain, (spl30_13 | ~ spl30_49 | ~ spl30_82), inference(avatar_contradiction_clause, [], [f5584])).
fof(f5584, plain, ($false | (spl30_13 | ~ spl30_49 | ~ spl30_82)), inference(subsumption_resolution, [], [f5574, f422])).
fof(f422, plain, (~ (e0 = op(e3, e0)) | spl30_13), inference(avatar_component_clause, [], [f421])).
fof(f5574, plain, ((e0 = op(e3, e0)) | (~ spl30_49 | ~ spl30_82)), inference(backward_demodulation, [], [f758, f576])).
fof(f5570, plain, (~ spl30_5 | ~ spl30_56 | spl30_191), inference(avatar_split_clause, [], [f5561, f1286, f603, f387])).
fof(f5561, plain, (~ (e0 = op(e3, e2)) | (~ spl30_56 | spl30_191)), inference(backward_demodulation, [], [f1288, f605])).
fof(f5569, plain, (~ spl30_49 | ~ spl30_56 | spl30_135), inference(avatar_split_clause, [], [f5560, f1015, f603, f574])).
fof(f1015, plain, (spl30_135 <=> (e0 = op(e0, op(e0, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl30_135])])).
fof(f5560, plain, (~ (e0 = op(e0, e3)) | (~ spl30_56 | spl30_135)), inference(backward_demodulation, [], [f1017, f605])).
fof(f1017, plain, (~ (e0 = op(e0, op(e0, e2))) | spl30_135), inference(avatar_component_clause, [], [f1015])).
fof(f5568, plain, (~ spl30_34 | ~ spl30_56 | spl30_134), inference(avatar_contradiction_clause, [], [f5567])).
fof(f5567, plain, ($false | (~ spl30_34 | ~ spl30_56 | spl30_134)), inference(subsumption_resolution, [], [f5559, f512])).
fof(f5559, plain, (~ (e1 = op(e1, e3)) | (~ spl30_56 | spl30_134)), inference(backward_demodulation, [], [f1012, f605])).
fof(f5565, plain, (~ spl30_52 | ~ spl30_56), inference(avatar_split_clause, [], [f5556, f603, f586])).
fof(f5556, plain, (~ (e3 = op(e0, e3)) | ~ spl30_56), inference(backward_demodulation, [], [f192, f605])).
fof(f5553, plain, (spl30_27 | ~ spl30_58 | ~ spl30_138), inference(avatar_contradiction_clause, [], [f5552])).
fof(f5552, plain, ($false | (spl30_27 | ~ spl30_58 | ~ spl30_138)), inference(subsumption_resolution, [], [f5541, f481])).
fof(f481, plain, (~ (e2 = op(e2, e1)) | spl30_27), inference(avatar_component_clause, [], [f480])).
fof(f5541, plain, ((e2 = op(e2, e1)) | (~ spl30_58 | ~ spl30_138)), inference(backward_demodulation, [], [f1031, f614])).
fof(f5549, plain, (spl30_46 | ~ spl30_58 | ~ spl30_125), inference(avatar_contradiction_clause, [], [f5548])).
fof(f5548, plain, ($false | (spl30_46 | ~ spl30_58 | ~ spl30_125)), inference(subsumption_resolution, [], [f5539, f562])).
fof(f562, plain, (~ (e1 = op(e1, e0)) | spl30_46), inference(avatar_component_clause, [], [f561])).
fof(f5539, plain, ((e1 = op(e1, e0)) | (~ spl30_58 | ~ spl30_125)), inference(backward_demodulation, [], [f968, f614])).
fof(f5547, plain, (~ spl30_10 | ~ spl30_58), inference(avatar_split_clause, [], [f5535, f612, f408])).
fof(f5535, plain, (~ (e1 = op(e3, e1)) | ~ spl30_58), inference(backward_demodulation, [], [f171, f614])).
fof(f5534, plain, (spl30_56 | ~ spl30_63 | ~ spl30_210), inference(avatar_split_clause, [], [f5530, f1382, f633, f603])).
fof(f5530, plain, ((e3 = op(e0, e2)) | (~ spl30_63 | ~ spl30_210)), inference(backward_demodulation, [], [f1383, f635])).
fof(f635, plain, ((op(e0, e0) = e2) | ~ spl30_63), inference(avatar_component_clause, [], [f633])).
fof(f5532, plain, (~ spl30_55 | ~ spl30_63), inference(avatar_split_clause, [], [f5522, f633, f599])).
fof(f5522, plain, (~ (e2 = op(e0, e2)) | ~ spl30_63), inference(backward_demodulation, [], [f188, f635])).
fof(f5531, plain, (~ spl30_31 | ~ spl30_63), inference(avatar_split_clause, [], [f5519, f633, f497])).
fof(f5519, plain, (~ (e2 = op(e2, e0)) | ~ spl30_63), inference(backward_demodulation, [], [f164, f635])).
fof(f5461, plain, (~ spl30_18 | ~ spl30_34), inference(avatar_split_clause, [], [f5460, f510, f442])).
fof(f5460, plain, (~ (e1 = op(e2, e3)) | ~ spl30_34), inference(forward_demodulation, [], [f184, f512])).
fof(f5458, plain, (~ spl30_17 | ~ spl30_22 | ~ spl30_28 | spl30_221), inference(avatar_split_clause, [], [f5380, f1437, f484, f459, f438])).
fof(f1437, plain, (spl30_221 <=> (e0 = op(e2, op(e2, op(e2, e2))))), introduced(avatar_definition, [new_symbols(naming, [spl30_221])])).
fof(f5380, plain, (~ (e0 = op(e2, e3)) | (~ spl30_22 | ~ spl30_28 | spl30_221)), inference(forward_demodulation, [], [f5379, f486])).
fof(f5379, plain, (~ (e0 = op(e2, op(e2, e1))) | (~ spl30_22 | spl30_221)), inference(forward_demodulation, [], [f1439, f461])).
fof(f1439, plain, (~ (e0 = op(e2, op(e2, op(e2, e2)))) | spl30_221), inference(avatar_component_clause, [], [f1437])).
fof(f5454, plain, (~ spl30_127 | spl30_78 | ~ spl30_82), inference(avatar_split_clause, [], [f5453, f756, f736, f975])).
fof(f5453, plain, (~ (e3 = op(e3, op(e0, e3))) | (spl30_78 | ~ spl30_82)), inference(forward_demodulation, [], [f738, f758])).
fof(f5452, plain, (~ spl30_130 | spl30_81 | ~ spl30_82), inference(avatar_split_clause, [], [f5451, f756, f751, f990])).
fof(f751, plain, (spl30_81 <=> (e0 = op(e0, op(e3, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl30_81])])).
fof(f5451, plain, (~ (e0 = op(e0, op(e0, e3))) | (spl30_81 | ~ spl30_82)), inference(forward_demodulation, [], [f753, f758])).
fof(f753, plain, (~ (e0 = op(e0, op(e3, e0))) | spl30_81), inference(avatar_component_clause, [], [f751])).
fof(f5440, plain, (~ spl30_11 | ~ spl30_28 | spl30_165), inference(avatar_split_clause, [], [f5416, f1161, f484, f412])).
fof(f1161, plain, (spl30_165 <=> (e2 = op(op(e2, e1), e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_165])])).
fof(f5416, plain, (~ (e2 = op(e3, e1)) | (~ spl30_28 | spl30_165)), inference(forward_demodulation, [], [f1163, f486])).
fof(f1163, plain, (~ (e2 = op(op(e2, e1), e1)) | spl30_165), inference(avatar_component_clause, [], [f1161])).
fof(f5396, plain, (~ spl30_55 | spl30_193), inference(avatar_contradiction_clause, [], [f5395])).
fof(f5395, plain, ($false | (~ spl30_55 | spl30_193)), inference(subsumption_resolution, [], [f5394, f601])).
fof(f5394, plain, (~ (e2 = op(e0, e2)) | (~ spl30_55 | spl30_193)), inference(forward_demodulation, [], [f1298, f601])).
fof(f1298, plain, (~ (e2 = op(e0, op(e0, e2))) | spl30_193), inference(avatar_component_clause, [], [f1296])).
fof(f5393, plain, (~ spl30_31 | spl30_169), inference(avatar_contradiction_clause, [], [f5392])).
fof(f5392, plain, ($false | (~ spl30_31 | spl30_169)), inference(subsumption_resolution, [], [f5391, f499])).
fof(f5391, plain, (~ (e2 = op(e2, e0)) | (~ spl30_31 | spl30_169)), inference(forward_demodulation, [], [f1182, f499])).
fof(f5388, plain, (spl30_21 | ~ spl30_31 | ~ spl30_170), inference(avatar_contradiction_clause, [], [f5387])).
fof(f5387, plain, ($false | (spl30_21 | ~ spl30_31 | ~ spl30_170)), inference(subsumption_resolution, [], [f5386, f456])).
fof(f456, plain, (~ (e0 = op(e2, e2)) | spl30_21), inference(avatar_component_clause, [], [f455])).
fof(f5386, plain, ((e0 = op(e2, e2)) | (~ spl30_31 | ~ spl30_170)), inference(forward_demodulation, [], [f1187, f499])).
fof(f5342, plain, (~ spl30_40 | ~ spl30_43 | spl30_216), inference(avatar_contradiction_clause, [], [f5341])).
fof(f5341, plain, ($false | (~ spl30_40 | ~ spl30_43 | spl30_216)), inference(subsumption_resolution, [], [f5333, f537])).
fof(f5333, plain, (~ (e3 = op(e1, e2)) | (~ spl30_43 | spl30_216)), inference(backward_demodulation, [], [f1414, f550])).
fof(f5337, plain, (~ spl30_23 | ~ spl30_43 | spl30_117), inference(avatar_split_clause, [], [f5327, f927, f548, f463])).
fof(f927, plain, (spl30_117 <=> (e2 = op(e2, op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl30_117])])).
fof(f5327, plain, (~ (e2 = op(e2, e2)) | (~ spl30_43 | spl30_117)), inference(backward_demodulation, [], [f929, f550])).
fof(f929, plain, (~ (e2 = op(e2, op(e1, e1))) | spl30_117), inference(avatar_component_clause, [], [f927])).
fof(f5334, plain, (~ spl30_11 | ~ spl30_43), inference(avatar_split_clause, [], [f5323, f548, f412])).
fof(f5323, plain, (~ (e2 = op(e3, e1)) | ~ spl30_43), inference(backward_demodulation, [], [f173, f550])).
fof(f5315, plain, (~ spl30_23 | ~ spl30_55), inference(avatar_split_clause, [], [f5306, f599, f463])).
fof(f5306, plain, (~ (e2 = op(e2, e2)) | ~ spl30_55), inference(backward_demodulation, [], [f176, f601])).
fof(f5302, plain, (~ spl30_57 | ~ spl30_62 | spl30_145), inference(avatar_contradiction_clause, [], [f5301])).
fof(f5301, plain, ($false | (~ spl30_57 | ~ spl30_62 | spl30_145)), inference(subsumption_resolution, [], [f5294, f610])).
fof(f5294, plain, (~ (e0 = op(e0, e1)) | (~ spl30_62 | spl30_145)), inference(backward_demodulation, [], [f1067, f631])).
fof(f1067, plain, (~ (e0 = op(e0, op(e0, e0))) | spl30_145), inference(avatar_component_clause, [], [f1065])).
fof(f5296, plain, (~ spl30_30 | ~ spl30_62), inference(avatar_split_clause, [], [f5287, f629, f493])).
fof(f5287, plain, (~ (e1 = op(e2, e0)) | ~ spl30_62), inference(backward_demodulation, [], [f164, f631])).
fof(f5278, plain, (spl30_34 | ~ spl30_40 | ~ spl30_113), inference(avatar_split_clause, [], [f5277, f907, f535, f510])).
fof(f5277, plain, ((e1 = op(e1, e3)) | (~ spl30_40 | ~ spl30_113)), inference(forward_demodulation, [], [f908, f537])).
fof(f908, plain, ((e1 = op(e1, op(e1, e2))) | ~ spl30_113), inference(avatar_component_clause, [], [f907])).
fof(f5274, plain, (spl30_31 | ~ spl30_57 | ~ spl30_138), inference(avatar_split_clause, [], [f5273, f1030, f608, f497])).
fof(f5273, plain, ((e2 = op(e2, e0)) | (~ spl30_57 | ~ spl30_138)), inference(forward_demodulation, [], [f1031, f610])).
fof(f5259, plain, (~ spl30_90 | ~ spl30_222), inference(avatar_contradiction_clause, [], [f5258])).
fof(f5258, plain, ($false | (~ spl30_90 | ~ spl30_222)), inference(subsumption_resolution, [], [f5257, f216])).
fof(f5257, plain, ((e2 = e3) | (~ spl30_90 | ~ spl30_222)), inference(forward_demodulation, [], [f1443, f796])).
fof(f1443, plain, ((e3 = op(e2, op(e2, e2))) | ~ spl30_222), inference(avatar_component_clause, [], [f1442])).
fof(f1442, plain, (spl30_222 <=> (e3 = op(e2, op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl30_222])])).
fof(f5242, plain, (~ spl30_64 | ~ spl30_52), inference(avatar_split_clause, [], [f5241, f586, f637])).
fof(f5241, plain, (~ (op(e0, e0) = e3) | ~ spl30_52), inference(forward_demodulation, [], [f189, f588])).
fof(f5240, plain, (~ spl30_62 | ~ spl30_57 | spl30_197), inference(avatar_split_clause, [], [f4740, f1315, f608, f629])).
fof(f4740, plain, (~ (op(e0, e0) = e1) | (~ spl30_57 | spl30_197)), inference(backward_demodulation, [], [f1317, f610])).
fof(f5232, plain, (~ spl30_56 | ~ spl30_17 | spl30_160), inference(avatar_split_clause, [], [f5127, f1136, f438, f603])).
fof(f5127, plain, (~ (e3 = op(e0, e2)) | (~ spl30_17 | spl30_160)), inference(backward_demodulation, [], [f1137, f440])).
fof(f1137, plain, (~ (e3 = op(op(e2, e3), e2)) | spl30_160), inference(avatar_component_clause, [], [f1136])).
fof(f5230, plain, (~ spl30_35 | ~ spl30_40 | spl30_179), inference(avatar_split_clause, [], [f5188, f1228, f535, f514])).
fof(f5188, plain, (~ (e2 = op(e1, e3)) | (~ spl30_40 | spl30_179)), inference(forward_demodulation, [], [f1230, f537])).
fof(f5229, plain, (~ spl30_31 | ~ spl30_45 | spl30_122), inference(avatar_split_clause, [], [f5067, f951, f557, f497])).
fof(f5067, plain, (~ (e2 = op(e2, e0)) | (~ spl30_45 | spl30_122)), inference(backward_demodulation, [], [f953, f559])).
fof(f559, plain, ((e0 = op(e1, e0)) | ~ spl30_45), inference(avatar_component_clause, [], [f557])).
fof(f5226, plain, (~ spl30_32 | ~ spl30_17 | spl30_161), inference(avatar_split_clause, [], [f5194, f1141, f438, f501])).
fof(f5194, plain, (~ (e3 = op(e2, e0)) | (~ spl30_17 | spl30_161)), inference(forward_demodulation, [], [f1143, f440])).
fof(f5173, plain, (~ spl30_2 | ~ spl30_11 | spl30_228), inference(avatar_contradiction_clause, [], [f5172])).
fof(f5172, plain, ($false | (~ spl30_2 | ~ spl30_11 | spl30_228)), inference(subsumption_resolution, [], [f5171, f414])).
fof(f5171, plain, (~ (e2 = op(e3, e1)) | (~ spl30_2 | spl30_228)), inference(forward_demodulation, [], [f1474, f376])).
fof(f376, plain, ((e1 = op(e3, e3)) | ~ spl30_2), inference(avatar_component_clause, [], [f374])).
fof(f1474, plain, (~ (e2 = op(e3, op(e3, e3))) | spl30_228), inference(avatar_component_clause, [], [f1472])).
fof(f1472, plain, (spl30_228 <=> (e2 = op(e3, op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl30_228])])).
fof(f5163, plain, (~ spl30_2 | spl30_27 | ~ spl30_201), inference(avatar_contradiction_clause, [], [f5162])).
fof(f5162, plain, ($false | (~ spl30_2 | spl30_27 | ~ spl30_201)), inference(subsumption_resolution, [], [f5158, f481])).
fof(f5158, plain, ((e2 = op(e2, e1)) | (~ spl30_2 | ~ spl30_201)), inference(backward_demodulation, [], [f1337, f376])).
fof(f5154, plain, (~ spl30_11 | ~ spl30_17 | ~ spl30_152), inference(avatar_contradiction_clause, [], [f5153])).
fof(f5153, plain, ($false | (~ spl30_11 | ~ spl30_17 | ~ spl30_152)), inference(subsumption_resolution, [], [f5152, f211])).
fof(f5152, plain, ((e0 = e1) | (~ spl30_11 | ~ spl30_17 | ~ spl30_152)), inference(forward_demodulation, [], [f5142, f440])).
fof(f5142, plain, ((e1 = op(e2, e3)) | (~ spl30_11 | ~ spl30_152)), inference(backward_demodulation, [], [f1100, f414])).
fof(f5151, plain, (~ spl30_11 | ~ spl30_28 | spl30_151), inference(avatar_contradiction_clause, [], [f5150])).
fof(f5150, plain, ($false | (~ spl30_11 | ~ spl30_28 | spl30_151)), inference(subsumption_resolution, [], [f5141, f486])).
fof(f5141, plain, (~ (e3 = op(e2, e1)) | (~ spl30_11 | spl30_151)), inference(backward_demodulation, [], [f1095, f414])).
fof(f1095, plain, (~ (e3 = op(op(e3, e1), e1)) | spl30_151), inference(avatar_component_clause, [], [f1093])).
fof(f5147, plain, (~ spl30_11 | ~ spl30_23 | spl30_73), inference(avatar_contradiction_clause, [], [f5146])).
fof(f5146, plain, ($false | (~ spl30_11 | ~ spl30_23 | spl30_73)), inference(subsumption_resolution, [], [f5138, f465])).
fof(f5138, plain, (~ (e2 = op(e2, e2)) | (~ spl30_11 | spl30_73)), inference(backward_demodulation, [], [f714, f414])).
fof(f714, plain, (~ (e2 = op(e2, op(e3, e1))) | spl30_73), inference(avatar_component_clause, [], [f712])).
fof(f5109, plain, (~ spl30_35 | ~ spl30_40 | spl30_175), inference(avatar_contradiction_clause, [], [f5108])).
fof(f5108, plain, ($false | (~ spl30_35 | ~ spl30_40 | spl30_175)), inference(subsumption_resolution, [], [f5099, f537])).
fof(f5099, plain, (~ (e3 = op(e1, e2)) | (~ spl30_35 | spl30_175)), inference(backward_demodulation, [], [f1211, f516])).
fof(f1211, plain, (~ (e3 = op(e1, op(e1, e3))) | spl30_175), inference(avatar_component_clause, [], [f1209])).
fof(f1209, plain, (spl30_175 <=> (e3 = op(e1, op(e1, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl30_175])])).
fof(f5107, plain, (~ spl30_18 | ~ spl30_35 | spl30_173), inference(avatar_split_clause, [], [f5097, f1199, f514, f442])).
fof(f1199, plain, (spl30_173 <=> (e1 = op(op(e1, e3), e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_173])])).
fof(f5097, plain, (~ (e1 = op(e2, e3)) | (~ spl30_35 | spl30_173)), inference(backward_demodulation, [], [f1201, f516])).
fof(f1201, plain, (~ (e1 = op(op(e1, e3), e3)) | spl30_173), inference(avatar_component_clause, [], [f1199])).
fof(f5106, plain, (~ spl30_23 | ~ spl30_35 | spl30_107), inference(avatar_contradiction_clause, [], [f5105])).
fof(f5105, plain, ($false | (~ spl30_23 | ~ spl30_35 | spl30_107)), inference(subsumption_resolution, [], [f5096, f465])).
fof(f5096, plain, (~ (e2 = op(e2, e2)) | (~ spl30_35 | spl30_107)), inference(backward_demodulation, [], [f879, f516])).
fof(f879, plain, (~ (e2 = op(e2, op(e1, e3))) | spl30_107), inference(avatar_component_clause, [], [f877])).
fof(f877, plain, (spl30_107 <=> (e2 = op(e2, op(e1, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl30_107])])).
fof(f5101, plain, (~ spl30_11 | ~ spl30_35 | spl30_76), inference(avatar_split_clause, [], [f5094, f727, f514, f412])).
fof(f5094, plain, (~ (e2 = op(e3, e1)) | (~ spl30_35 | spl30_76)), inference(backward_demodulation, [], [f728, f516])).
fof(f5089, plain, (~ spl30_42 | ~ spl30_213), inference(avatar_contradiction_clause, [], [f5088])).
fof(f5088, plain, ($false | (~ spl30_42 | ~ spl30_213)), inference(subsumption_resolution, [], [f5087, f214])).
fof(f5087, plain, ((e1 = e2) | (~ spl30_42 | ~ spl30_213)), inference(forward_demodulation, [], [f5086, f546])).
fof(f5086, plain, ((e2 = op(e1, e1)) | (~ spl30_42 | ~ spl30_213)), inference(forward_demodulation, [], [f5079, f546])).
fof(f5079, plain, ((e2 = op(e1, op(e1, e1))) | (~ spl30_42 | ~ spl30_213)), inference(backward_demodulation, [], [f1398, f546])).
fof(f1398, plain, ((e2 = op(e1, op(e1, op(e1, e1)))) | ~ spl30_213), inference(avatar_component_clause, [], [f1397])).
fof(f5071, plain, (~ spl30_33 | ~ spl30_45), inference(avatar_split_clause, [], [f5063, f557, f506])).
fof(f5063, plain, (~ (e0 = op(e1, e3)) | ~ spl30_45), inference(backward_demodulation, [], [f195, f559])).
fof(f5070, plain, (~ spl30_29 | ~ spl30_45), inference(avatar_split_clause, [], [f5061, f557, f489])).
fof(f5061, plain, (~ (e0 = op(e2, e0)) | ~ spl30_45), inference(backward_demodulation, [], [f166, f559])).
fof(f5060, plain, (spl30_17 | ~ spl30_5 | ~ spl30_70), inference(avatar_split_clause, [], [f5059, f698, f387, f438])).
fof(f5059, plain, ((e0 = op(e2, e3)) | (~ spl30_5 | ~ spl30_70)), inference(forward_demodulation, [], [f700, f389])).
fof(f5026, plain, (~ spl30_4 | ~ spl30_28 | spl30_94), inference(avatar_split_clause, [], [f4373, f814, f484, f382])).
fof(f4373, plain, (~ (e3 = op(e3, e3)) | (~ spl30_28 | spl30_94)), inference(forward_demodulation, [], [f816, f486])).
fof(f4980, plain, (~ spl30_33 | ~ spl30_40 | ~ spl30_43 | spl30_215), inference(avatar_split_clause, [], [f4973, f1407, f548, f535, f506])).
fof(f4973, plain, (~ (e0 = op(e1, e3)) | (~ spl30_40 | ~ spl30_43 | spl30_215)), inference(backward_demodulation, [], [f4828, f537])).
fof(f4828, plain, (~ (e0 = op(e1, op(e1, e2))) | (~ spl30_43 | spl30_215)), inference(forward_demodulation, [], [f1409, f550])).
fof(f4979, plain, (spl30_35 | ~ spl30_40 | ~ spl30_179), inference(avatar_contradiction_clause, [], [f4978])).
fof(f4978, plain, ($false | (spl30_35 | ~ spl30_40 | ~ spl30_179)), inference(subsumption_resolution, [], [f4972, f515])).
fof(f4972, plain, ((e2 = op(e1, e3)) | (~ spl30_40 | ~ spl30_179)), inference(backward_demodulation, [], [f1229, f537])).
fof(f1229, plain, ((e2 = op(e1, op(e1, e2))) | ~ spl30_179), inference(avatar_component_clause, [], [f1228])).
fof(f4964, plain, (spl30_41 | ~ spl30_46 | ~ spl30_184), inference(avatar_contradiction_clause, [], [f4963])).
fof(f4963, plain, ($false | (spl30_41 | ~ spl30_46 | ~ spl30_184)), inference(subsumption_resolution, [], [f4957, f541])).
fof(f541, plain, (~ (e0 = op(e1, e1)) | spl30_41), inference(avatar_component_clause, [], [f540])).
fof(f4957, plain, ((e0 = op(e1, e1)) | (~ spl30_46 | ~ spl30_184)), inference(backward_demodulation, [], [f1255, f563])).
fof(f4962, plain, (~ spl30_46 | spl30_183), inference(avatar_contradiction_clause, [], [f4961])).
fof(f4961, plain, ($false | (~ spl30_46 | spl30_183)), inference(subsumption_resolution, [], [f4956, f563])).
fof(f4956, plain, (~ (e1 = op(e1, e0)) | (~ spl30_46 | spl30_183)), inference(backward_demodulation, [], [f1250, f563])).
fof(f4916, plain, (spl30_52 | ~ spl30_16 | ~ spl30_82), inference(avatar_split_clause, [], [f4915, f756, f433, f586])).
fof(f4915, plain, ((e3 = op(e0, e3)) | (~ spl30_16 | ~ spl30_82)), inference(forward_demodulation, [], [f758, f435])).
fof(f4914, plain, (spl30_40 | ~ spl30_28 | ~ spl30_98), inference(avatar_split_clause, [], [f4913, f834, f484, f535])).
fof(f4913, plain, ((e3 = op(e1, e2)) | (~ spl30_28 | ~ spl30_98)), inference(forward_demodulation, [], [f836, f486])).
fof(f4912, plain, (~ spl30_45 | ~ spl30_57 | spl30_125), inference(avatar_split_clause, [], [f4911, f966, f608, f557])).
fof(f4911, plain, (~ (e0 = op(e1, e0)) | (~ spl30_57 | spl30_125)), inference(forward_demodulation, [], [f967, f610])).
fof(f4902, plain, (spl30_54 | ~ spl30_63 | ~ spl30_205), inference(avatar_split_clause, [], [f4901, f1358, f633, f595])).
fof(f4901, plain, ((e1 = op(e0, e2)) | (~ spl30_63 | ~ spl30_205)), inference(forward_demodulation, [], [f1359, f635])).
fof(f4895, plain, (~ spl30_56 | ~ spl30_63 | spl30_210), inference(avatar_split_clause, [], [f4820, f1382, f633, f603])).
fof(f4820, plain, (~ (e3 = op(e0, e2)) | (~ spl30_63 | spl30_210)), inference(forward_demodulation, [], [f1384, f635])).
fof(f4893, plain, (~ spl30_51 | ~ spl30_63), inference(avatar_split_clause, [], [f4892, f633, f582])).
fof(f4892, plain, (~ (e2 = op(e0, e3)) | ~ spl30_63), inference(forward_demodulation, [], [f189, f635])).
fof(f4885, plain, (~ spl30_39 | ~ spl30_43), inference(avatar_split_clause, [], [f4884, f548, f531])).
fof(f4884, plain, (~ (e2 = op(e1, e2)) | ~ spl30_43), inference(forward_demodulation, [], [f196, f550])).
fof(f4881, plain, (spl30_38 | ~ spl30_63 | ~ spl30_144), inference(avatar_split_clause, [], [f4843, f1060, f633, f527])).
fof(f4843, plain, ((e1 = op(e1, e2)) | (~ spl30_63 | ~ spl30_144)), inference(forward_demodulation, [], [f1061, f635])).
fof(f1061, plain, ((e1 = op(e1, op(e0, e0))) | ~ spl30_144), inference(avatar_component_clause, [], [f1060])).
fof(f4870, plain, (~ spl30_31 | ~ spl30_57 | spl30_138), inference(avatar_split_clause, [], [f4848, f1030, f608, f497])).
fof(f4848, plain, (~ (e2 = op(e2, e0)) | (~ spl30_57 | spl30_138)), inference(forward_demodulation, [], [f1032, f610])).
fof(f4868, plain, (~ spl30_19 | ~ spl30_28 | spl30_95), inference(avatar_split_clause, [], [f4271, f819, f484, f446])).
fof(f819, plain, (spl30_95 <=> (e2 = op(e2, op(e2, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl30_95])])).
fof(f4271, plain, (~ (e2 = op(e2, e3)) | (~ spl30_28 | spl30_95)), inference(forward_demodulation, [], [f821, f486])).
fof(f821, plain, (~ (e2 = op(e2, op(e2, e1))) | spl30_95), inference(avatar_component_clause, [], [f819])).
fof(f4867, plain, (~ spl30_18 | ~ spl30_28 | spl30_167), inference(avatar_split_clause, [], [f4298, f1171, f484, f442])).
fof(f4298, plain, (~ (e1 = op(e2, e3)) | (~ spl30_28 | spl30_167)), inference(forward_demodulation, [], [f1173, f486])).
fof(f4819, plain, (~ spl30_50 | ~ spl30_56 | ~ spl30_63 | spl30_209), inference(avatar_contradiction_clause, [], [f4818])).
fof(f4818, plain, ($false | (~ spl30_50 | ~ spl30_56 | ~ spl30_63 | spl30_209)), inference(subsumption_resolution, [], [f4817, f580])).
fof(f4817, plain, (~ (e1 = op(e0, e3)) | (~ spl30_56 | ~ spl30_63 | spl30_209)), inference(forward_demodulation, [], [f4816, f605])).
fof(f4816, plain, (~ (e1 = op(e0, op(e0, e2))) | (~ spl30_63 | spl30_209)), inference(forward_demodulation, [], [f1379, f635])).
fof(f4810, plain, (~ spl30_5 | spl30_61 | ~ spl30_69), inference(avatar_contradiction_clause, [], [f4809])).
fof(f4809, plain, ($false | (~ spl30_5 | spl30_61 | ~ spl30_69)), inference(subsumption_resolution, [], [f4801, f626])).
fof(f626, plain, (~ (e0 = op(e0, e0)) | spl30_61), inference(avatar_component_clause, [], [f625])).
fof(f4801, plain, ((e0 = op(e0, e0)) | (~ spl30_5 | ~ spl30_69)), inference(backward_demodulation, [], [f694, f389])).
fof(f694, plain, ((e0 = op(e0, op(e3, e2))) | ~ spl30_69), inference(avatar_component_clause, [], [f693])).
fof(f693, plain, (spl30_69 <=> (e0 = op(e0, op(e3, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl30_69])])).
fof(f4792, plain, (~ spl30_8 | ~ spl30_16), inference(avatar_split_clause, [], [f4785, f433, f399])).
fof(f4785, plain, (~ (e3 = op(e3, e2)) | ~ spl30_16), inference(backward_demodulation, [], [f206, f435])).
fof(f4781, plain, (~ spl30_43 | spl30_53 | ~ spl30_119), inference(avatar_contradiction_clause, [], [f4780])).
fof(f4780, plain, ($false | (~ spl30_43 | spl30_53 | ~ spl30_119)), inference(subsumption_resolution, [], [f4773, f592])).
fof(f4773, plain, ((e0 = op(e0, e2)) | (~ spl30_43 | ~ spl30_119)), inference(backward_demodulation, [], [f938, f550])).
fof(f938, plain, ((e0 = op(e0, op(e1, e1))) | ~ spl30_119), inference(avatar_component_clause, [], [f937])).
fof(f937, plain, (spl30_119 <=> (e0 = op(e0, op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl30_119])])).
fof(f4779, plain, (~ spl30_38 | ~ spl30_43 | spl30_118), inference(avatar_contradiction_clause, [], [f4778])).
fof(f4778, plain, ($false | (~ spl30_38 | ~ spl30_43 | spl30_118)), inference(subsumption_resolution, [], [f4772, f529])).
fof(f4772, plain, (~ (e1 = op(e1, e2)) | (~ spl30_43 | spl30_118)), inference(backward_demodulation, [], [f934, f550])).
fof(f4750, plain, (~ spl30_57 | spl30_195), inference(avatar_contradiction_clause, [], [f4749])).
fof(f4749, plain, ($false | (~ spl30_57 | spl30_195)), inference(subsumption_resolution, [], [f4738, f610])).
fof(f4738, plain, (~ (e0 = op(e0, e1)) | (~ spl30_57 | spl30_195)), inference(backward_demodulation, [], [f1307, f610])).
fof(f4729, plain, (~ spl30_53 | ~ spl30_63 | spl30_145), inference(avatar_split_clause, [], [f4720, f1065, f633, f591])).
fof(f4720, plain, (~ (e0 = op(e0, e2)) | (~ spl30_63 | spl30_145)), inference(backward_demodulation, [], [f1067, f635])).
fof(f4726, plain, (~ spl30_23 | ~ spl30_63 | spl30_143), inference(avatar_contradiction_clause, [], [f4725])).
fof(f4725, plain, ($false | (~ spl30_23 | ~ spl30_63 | spl30_143)), inference(subsumption_resolution, [], [f4718, f465])).
fof(f4718, plain, (~ (e2 = op(e2, e2)) | (~ spl30_63 | spl30_143)), inference(backward_demodulation, [], [f1057, f635])).
fof(f1057, plain, (~ (e2 = op(e2, op(e0, e0))) | spl30_143), inference(avatar_component_clause, [], [f1055])).
fof(f1055, plain, (spl30_143 <=> (e2 = op(e2, op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl30_143])])).
fof(f4723, plain, (~ spl30_59 | ~ spl30_63), inference(avatar_split_clause, [], [f4714, f633, f616])).
fof(f4714, plain, (~ (e2 = op(e0, e1)) | ~ spl30_63), inference(backward_demodulation, [], [f187, f635])).
fof(f4695, plain, (~ spl30_53 | ~ spl30_3 | spl30_203), inference(avatar_split_clause, [], [f4694, f1346, f378, f591])).
fof(f4694, plain, (~ (e0 = op(e0, e2)) | (~ spl30_3 | spl30_203)), inference(forward_demodulation, [], [f1348, f380])).
fof(f1348, plain, (~ (e0 = op(e0, op(e3, e3))) | spl30_203), inference(avatar_component_clause, [], [f1346])).
fof(f4681, plain, (~ spl30_59 | ~ spl30_205 | spl30_207), inference(avatar_split_clause, [], [f4580, f1367, f1358, f616])).
fof(f4580, plain, (~ (e2 = op(e0, e1)) | (~ spl30_205 | spl30_207)), inference(backward_demodulation, [], [f1369, f1359])).
fof(f4680, plain, (~ spl30_60 | ~ spl30_205 | spl30_206), inference(avatar_split_clause, [], [f4583, f1362, f1358, f620])).
fof(f1362, plain, (spl30_206 <=> (e3 = op(e0, op(e0, op(e0, e0))))), introduced(avatar_definition, [new_symbols(naming, [spl30_206])])).
fof(f4583, plain, (~ (e3 = op(e0, e1)) | (~ spl30_205 | spl30_206)), inference(backward_demodulation, [], [f1364, f1359])).
fof(f1364, plain, (~ (e3 = op(e0, op(e0, op(e0, e0)))) | spl30_206), inference(avatar_component_clause, [], [f1362])).
fof(f4678, plain, (~ spl30_57 | ~ spl30_38 | spl30_114), inference(avatar_split_clause, [], [f4418, f912, f527, f608])).
fof(f912, plain, (spl30_114 <=> (e0 = op(e0, op(e1, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl30_114])])).
fof(f4418, plain, (~ (e0 = op(e0, e1)) | (~ spl30_38 | spl30_114)), inference(backward_demodulation, [], [f914, f529])).
fof(f914, plain, (~ (e0 = op(e0, op(e1, e2))) | spl30_114), inference(avatar_component_clause, [], [f912])).
fof(f4677, plain, (~ spl30_57 | ~ spl30_10 | spl30_75), inference(avatar_split_clause, [], [f4459, f722, f408, f608])).
fof(f722, plain, (spl30_75 <=> (e0 = op(e0, op(e3, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl30_75])])).
fof(f4459, plain, (~ (e0 = op(e0, e1)) | (~ spl30_10 | spl30_75)), inference(forward_demodulation, [], [f724, f410])).
fof(f724, plain, (~ (e0 = op(e0, op(e3, e1))) | spl30_75), inference(avatar_component_clause, [], [f722])).
fof(f4656, plain, (~ spl30_43 | ~ spl30_38 | spl30_179), inference(avatar_split_clause, [], [f4475, f1228, f527, f548])).
fof(f4475, plain, (~ (e2 = op(e1, e1)) | (~ spl30_38 | spl30_179)), inference(forward_demodulation, [], [f1230, f529])).
fof(f4642, plain, (~ spl30_13 | spl30_157), inference(avatar_contradiction_clause, [], [f4641])).
fof(f4641, plain, ($false | (~ spl30_13 | spl30_157)), inference(subsumption_resolution, [], [f4633, f423])).
fof(f4633, plain, (~ (e0 = op(e3, e0)) | (~ spl30_13 | spl30_157)), inference(backward_demodulation, [], [f1124, f423])).
fof(f4634, plain, (~ spl30_5 | ~ spl30_13), inference(avatar_split_clause, [], [f4623, f421, f387])).
fof(f4623, plain, (~ (e0 = op(e3, e2)) | ~ spl30_13), inference(backward_demodulation, [], [f206, f423])).
fof(f4609, plain, (~ spl30_56 | ~ spl30_64), inference(avatar_split_clause, [], [f4602, f637, f603])).
fof(f4602, plain, (~ (e3 = op(e0, e2)) | ~ spl30_64), inference(backward_demodulation, [], [f188, f639])).
fof(f4591, plain, (spl30_8 | ~ spl30_3 | ~ spl30_200), inference(avatar_split_clause, [], [f4590, f1331, f378, f399])).
fof(f1331, plain, (spl30_200 <=> (e3 = op(e3, op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl30_200])])).
fof(f4590, plain, ((e3 = op(e3, e2)) | (~ spl30_3 | ~ spl30_200)), inference(forward_demodulation, [], [f1332, f380])).
fof(f1332, plain, ((e3 = op(e3, op(e3, e3))) | ~ spl30_200), inference(avatar_component_clause, [], [f1331])).
fof(f4589, plain, (spl30_53 | ~ spl30_3 | ~ spl30_203), inference(avatar_split_clause, [], [f4588, f1346, f378, f591])).
fof(f4588, plain, ((e0 = op(e0, e2)) | (~ spl30_3 | ~ spl30_203)), inference(forward_demodulation, [], [f1347, f380])).
fof(f4572, plain, (~ spl30_62 | ~ spl30_50), inference(avatar_split_clause, [], [f4571, f578, f629])).
fof(f4571, plain, (~ (op(e0, e0) = e1) | ~ spl30_50), inference(forward_demodulation, [], [f189, f580])).
fof(f4559, plain, (~ spl30_7 | ~ spl30_3), inference(avatar_split_clause, [], [f4558, f378, f395])).
fof(f4558, plain, (~ (e2 = op(e3, e2)) | ~ spl30_3), inference(forward_demodulation, [], [f210, f380])).
fof(f210, plain, ~ (op(e3, e2) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f4557, plain, (~ spl30_6 | ~ spl30_10), inference(avatar_split_clause, [], [f4556, f408, f391])).
fof(f4556, plain, (~ (e1 = op(e3, e2)) | ~ spl30_10), inference(forward_demodulation, [], [f208, f410])).
fof(f4553, plain, (~ spl30_6 | ~ spl30_38), inference(avatar_split_clause, [], [f4552, f527, f391])).
fof(f4552, plain, (~ (e1 = op(e3, e2)) | ~ spl30_38), inference(forward_demodulation, [], [f179, f529])).
fof(f4551, plain, (~ spl30_6 | ~ spl30_28 | spl30_166), inference(avatar_split_clause, [], [f4208, f1166, f484, f391])).
fof(f4208, plain, (~ (e1 = op(e3, e2)) | (~ spl30_28 | spl30_166)), inference(backward_demodulation, [], [f1167, f486])).
fof(f4550, plain, (~ spl30_6 | ~ spl30_3 | spl30_226), inference(avatar_split_clause, [], [f4259, f1462, f378, f391])).
fof(f4259, plain, (~ (e1 = op(e3, e2)) | (~ spl30_3 | spl30_226)), inference(backward_demodulation, [], [f1464, f380])).
fof(f4544, plain, (~ spl30_23 | spl30_90), inference(avatar_contradiction_clause, [], [f4543])).
fof(f4543, plain, ($false | (~ spl30_23 | spl30_90)), inference(subsumption_resolution, [], [f4542, f465])).
fof(f4542, plain, (~ (e2 = op(e2, e2)) | (~ spl30_23 | spl30_90)), inference(forward_demodulation, [], [f797, f465])).
fof(f4529, plain, (~ spl30_41 | ~ spl30_61 | spl30_119), inference(avatar_contradiction_clause, [], [f4528])).
fof(f4528, plain, ($false | (~ spl30_41 | ~ spl30_61 | spl30_119)), inference(subsumption_resolution, [], [f4527, f627])).
fof(f4527, plain, (~ (e0 = op(e0, e0)) | (~ spl30_41 | spl30_119)), inference(forward_demodulation, [], [f939, f542])).
fof(f939, plain, (~ (e0 = op(e0, op(e1, e1))) | spl30_119), inference(avatar_component_clause, [], [f937])).
fof(f4448, plain, (spl30_2 | ~ spl30_36 | ~ spl30_76 | ~ spl30_153), inference(avatar_contradiction_clause, [], [f4447])).
fof(f4447, plain, ($false | (spl30_2 | ~ spl30_36 | ~ spl30_76 | ~ spl30_153)), inference(subsumption_resolution, [], [f4437, f375])).
fof(f375, plain, (~ (e1 = op(e3, e3)) | spl30_2), inference(avatar_component_clause, [], [f374])).
fof(f4437, plain, ((e1 = op(e3, e3)) | (~ spl30_36 | ~ spl30_76 | ~ spl30_153)), inference(backward_demodulation, [], [f4130, f520])).
fof(f4130, plain, ((e1 = op(e3, op(e1, e3))) | (~ spl30_76 | ~ spl30_153)), inference(backward_demodulation, [], [f1104, f729])).
fof(f1104, plain, ((e1 = op(e3, op(e3, e1))) | ~ spl30_153), inference(avatar_component_clause, [], [f1103])).
fof(f4446, plain, (~ spl30_36 | spl30_175), inference(avatar_contradiction_clause, [], [f4445])).
fof(f4445, plain, ($false | (~ spl30_36 | spl30_175)), inference(subsumption_resolution, [], [f4436, f520])).
fof(f4436, plain, (~ (e3 = op(e1, e3)) | (~ spl30_36 | spl30_175)), inference(backward_demodulation, [], [f1211, f520])).
fof(f4429, plain, (~ spl30_38 | spl30_177), inference(avatar_contradiction_clause, [], [f4428])).
fof(f4428, plain, ($false | (~ spl30_38 | spl30_177)), inference(subsumption_resolution, [], [f4419, f529])).
fof(f4419, plain, (~ (e1 = op(e1, e2)) | (~ spl30_38 | spl30_177)), inference(backward_demodulation, [], [f1220, f529])).
fof(f4408, plain, (~ spl30_47 | ~ spl30_50 | ~ spl30_188), inference(avatar_contradiction_clause, [], [f4407])).
fof(f4407, plain, ($false | (~ spl30_47 | ~ spl30_50 | ~ spl30_188)), inference(subsumption_resolution, [], [f4406, f216])).
fof(f4406, plain, ((e2 = e3) | (~ spl30_47 | ~ spl30_50 | ~ spl30_188)), inference(forward_demodulation, [], [f4399, f567])).
fof(f4399, plain, ((e3 = op(e1, e0)) | (~ spl30_50 | ~ spl30_188)), inference(backward_demodulation, [], [f1274, f580])).
fof(f4390, plain, (spl30_51 | ~ spl30_56 | ~ spl30_193), inference(avatar_contradiction_clause, [], [f4389])).
fof(f4389, plain, ($false | (spl30_51 | ~ spl30_56 | ~ spl30_193)), inference(subsumption_resolution, [], [f4381, f583])).
fof(f583, plain, (~ (e2 = op(e0, e3)) | spl30_51), inference(avatar_component_clause, [], [f582])).
fof(f4381, plain, ((e2 = op(e0, e3)) | (~ spl30_56 | ~ spl30_193)), inference(backward_demodulation, [], [f1297, f605])).
fof(f4388, plain, (spl30_15 | ~ spl30_56 | ~ spl30_192), inference(avatar_contradiction_clause, [], [f4387])).
fof(f4387, plain, ($false | (spl30_15 | ~ spl30_56 | ~ spl30_192)), inference(subsumption_resolution, [], [f4380, f430])).
fof(f430, plain, (~ (e2 = op(e3, e0)) | spl30_15), inference(avatar_component_clause, [], [f429])).
fof(f4380, plain, ((e2 = op(e3, e0)) | (~ spl30_56 | ~ spl30_192)), inference(backward_demodulation, [], [f1293, f605])).
fof(f4370, plain, (spl30_56 | ~ spl30_5 | ~ spl30_147), inference(avatar_split_clause, [], [f4369, f1074, f387, f603])).
fof(f4369, plain, ((e3 = op(e0, e2)) | (~ spl30_5 | ~ spl30_147)), inference(forward_demodulation, [], [f1075, f389])).
fof(f1075, plain, ((e3 = op(op(e3, e2), e2)) | ~ spl30_147), inference(avatar_component_clause, [], [f1074])).
fof(f4363, plain, (spl30_54 | ~ spl30_30 | ~ spl30_104), inference(avatar_split_clause, [], [f4362, f863, f493, f595])).
fof(f4362, plain, ((e1 = op(e0, e2)) | (~ spl30_30 | ~ spl30_104)), inference(forward_demodulation, [], [f865, f495])).
fof(f4359, plain, (~ spl30_56 | ~ spl30_16 | ~ spl30_104), inference(avatar_split_clause, [], [f4358, f863, f433, f603])).
fof(f4358, plain, (~ (e3 = op(e0, e2)) | (~ spl30_16 | ~ spl30_104)), inference(forward_demodulation, [], [f3936, f435])).
fof(f3936, plain, (~ (op(e0, e2) = op(e3, e0)) | ~ spl30_104), inference(backward_demodulation, [], [f168, f865])).
fof(f4357, plain, (~ spl30_56 | ~ spl30_28 | ~ spl30_104), inference(avatar_split_clause, [], [f4356, f863, f484, f603])).
fof(f4356, plain, (~ (e3 = op(e0, e2)) | (~ spl30_28 | ~ spl30_104)), inference(forward_demodulation, [], [f3937, f486])).
fof(f3937, plain, (~ (op(e0, e2) = op(e2, e1)) | ~ spl30_104), inference(backward_demodulation, [], [f199, f865])).
fof(f4346, plain, (~ spl30_40 | ~ spl30_41 | ~ spl30_47 | spl30_212), inference(avatar_split_clause, [], [f4288, f1392, f565, f540, f535])).
fof(f4288, plain, (~ (e3 = op(e1, e2)) | (~ spl30_41 | ~ spl30_47 | spl30_212)), inference(forward_demodulation, [], [f4287, f567])).
fof(f4287, plain, (~ (e3 = op(e1, op(e1, e0))) | (~ spl30_41 | spl30_212)), inference(forward_demodulation, [], [f1394, f542])).
fof(f4342, plain, (spl30_34 | ~ spl30_10 | ~ spl30_76), inference(avatar_split_clause, [], [f4341, f727, f408, f510])).
fof(f4341, plain, ((e1 = op(e1, e3)) | (~ spl30_10 | ~ spl30_76)), inference(forward_demodulation, [], [f729, f410])).
fof(f4338, plain, (~ spl30_36 | ~ spl30_16 | ~ spl30_76), inference(avatar_split_clause, [], [f4337, f727, f433, f518])).
fof(f4337, plain, (~ (e3 = op(e1, e3)) | (~ spl30_16 | ~ spl30_76)), inference(forward_demodulation, [], [f4123, f435])).
fof(f4123, plain, (~ (op(e1, e3) = op(e3, e0)) | ~ spl30_76), inference(backward_demodulation, [], [f205, f729])).
fof(f4323, plain, (~ spl30_3 | ~ spl30_28 | ~ spl30_94), inference(avatar_contradiction_clause, [], [f4322])).
fof(f4322, plain, ($false | (~ spl30_3 | ~ spl30_28 | ~ spl30_94)), inference(subsumption_resolution, [], [f4321, f216])).
fof(f4321, plain, ((e2 = e3) | (~ spl30_3 | ~ spl30_28 | ~ spl30_94)), inference(forward_demodulation, [], [f4320, f380])).
fof(f4320, plain, ((e3 = op(e3, e3)) | (~ spl30_28 | ~ spl30_94)), inference(forward_demodulation, [], [f815, f486])).
fof(f4305, plain, (~ spl30_17 | ~ spl30_54 | ~ spl30_160), inference(avatar_contradiction_clause, [], [f4304])).
fof(f4304, plain, ($false | (~ spl30_17 | ~ spl30_54 | ~ spl30_160)), inference(subsumption_resolution, [], [f4303, f215])).
fof(f4303, plain, ((e1 = e3) | (~ spl30_17 | ~ spl30_54 | ~ spl30_160)), inference(forward_demodulation, [], [f4302, f597])).
fof(f4302, plain, ((e3 = op(e0, e2)) | (~ spl30_17 | ~ spl30_160)), inference(forward_demodulation, [], [f1138, f440])).
fof(f4264, plain, (~ spl30_3 | ~ spl30_23 | spl30_201), inference(avatar_contradiction_clause, [], [f4263])).
fof(f4263, plain, ($false | (~ spl30_3 | ~ spl30_23 | spl30_201)), inference(subsumption_resolution, [], [f4255, f465])).
fof(f4255, plain, (~ (e2 = op(e2, e2)) | (~ spl30_3 | spl30_201)), inference(backward_demodulation, [], [f1338, f380])).
fof(f4247, plain, (~ spl30_17 | spl30_51 | ~ spl30_159), inference(avatar_contradiction_clause, [], [f4246])).
fof(f4246, plain, ($false | (~ spl30_17 | spl30_51 | ~ spl30_159)), inference(subsumption_resolution, [], [f4241, f583])).
fof(f4241, plain, ((e2 = op(e0, e3)) | (~ spl30_17 | ~ spl30_159)), inference(backward_demodulation, [], [f1132, f440])).
fof(f1132, plain, ((e2 = op(op(e2, e3), e3)) | ~ spl30_159), inference(avatar_component_clause, [], [f1131])).
fof(f1131, plain, (spl30_159 <=> (e2 = op(op(e2, e3), e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_159])])).
fof(f4245, plain, (~ spl30_17 | ~ spl30_61 | spl30_87), inference(avatar_contradiction_clause, [], [f4244])).
fof(f4244, plain, ($false | (~ spl30_17 | ~ spl30_61 | spl30_87)), inference(subsumption_resolution, [], [f4240, f627])).
fof(f4240, plain, (~ (e0 = op(e0, e0)) | (~ spl30_17 | spl30_87)), inference(backward_demodulation, [], [f782, f440])).
fof(f4242, plain, (~ spl30_16 | ~ spl30_17 | spl30_84), inference(avatar_split_clause, [], [f4237, f765, f438, f433])).
fof(f4237, plain, (~ (e3 = op(e3, e0)) | (~ spl30_17 | spl30_84)), inference(backward_demodulation, [], [f767, f440])).
fof(f4229, plain, (~ spl30_23 | ~ spl30_218), inference(avatar_contradiction_clause, [], [f4228])).
fof(f4228, plain, ($false | (~ spl30_23 | ~ spl30_218)), inference(subsumption_resolution, [], [f4227, f216])).
fof(f4227, plain, ((e2 = e3) | (~ spl30_23 | ~ spl30_218)), inference(forward_demodulation, [], [f4226, f465])).
fof(f4226, plain, ((e3 = op(e2, e2)) | (~ spl30_23 | ~ spl30_218)), inference(forward_demodulation, [], [f4221, f465])).
fof(f4221, plain, ((e3 = op(e2, op(e2, e2))) | (~ spl30_23 | ~ spl30_218)), inference(backward_demodulation, [], [f1423, f465])).
fof(f1423, plain, ((e3 = op(e2, op(e2, op(e2, e2)))) | ~ spl30_218), inference(avatar_component_clause, [], [f1422])).
fof(f4217, plain, (spl30_18 | ~ spl30_28 | ~ spl30_167), inference(avatar_contradiction_clause, [], [f4216])).
fof(f4216, plain, ($false | (spl30_18 | ~ spl30_28 | ~ spl30_167)), inference(subsumption_resolution, [], [f4209, f443])).
fof(f4209, plain, ((e1 = op(e2, e3)) | (~ spl30_28 | ~ spl30_167)), inference(backward_demodulation, [], [f1172, f486])).
fof(f4175, plain, (~ spl30_19 | ~ spl30_40 | spl30_112), inference(avatar_split_clause, [], [f4168, f902, f535, f446])).
fof(f4168, plain, (~ (e2 = op(e2, e3)) | (~ spl30_40 | spl30_112)), inference(backward_demodulation, [], [f904, f537])).
fof(f4163, plain, (~ spl30_41 | ~ spl30_47 | spl30_214), inference(avatar_contradiction_clause, [], [f4162])).
fof(f4162, plain, ($false | (~ spl30_41 | ~ spl30_47 | spl30_214)), inference(subsumption_resolution, [], [f4156, f567])).
fof(f4156, plain, (~ (e2 = op(e1, e0)) | (~ spl30_41 | spl30_214)), inference(backward_demodulation, [], [f1404, f542])).
fof(f1404, plain, (~ (e2 = op(e1, op(e1, e1))) | spl30_214), inference(avatar_component_clause, [], [f1402])).
fof(f4121, plain, (spl30_16 | ~ spl30_52 | ~ spl30_82), inference(avatar_split_clause, [], [f4120, f756, f586, f433])).
fof(f4120, plain, ((e3 = op(e3, e0)) | (~ spl30_52 | ~ spl30_82)), inference(forward_demodulation, [], [f758, f588])).
fof(f4111, plain, (~ spl30_25 | ~ spl30_30 | spl30_171), inference(avatar_split_clause, [], [f4110, f1190, f493, f472])).
fof(f4110, plain, (~ (e0 = op(e2, e1)) | (~ spl30_30 | spl30_171)), inference(forward_demodulation, [], [f1192, f495])).
fof(f4105, plain, (~ spl30_25 | ~ spl30_59 | spl30_195), inference(avatar_split_clause, [], [f4104, f1305, f616, f472])).
fof(f4104, plain, (~ (e0 = op(e2, e1)) | (~ spl30_59 | spl30_195)), inference(forward_demodulation, [], [f1307, f618])).
fof(f4095, plain, (~ spl30_20 | ~ spl30_52), inference(avatar_split_clause, [], [f4094, f586, f450])).
fof(f4094, plain, (~ (e3 = op(e2, e3)) | ~ spl30_52), inference(forward_demodulation, [], [f182, f588])).
fof(f4091, plain, (~ spl30_16 | ~ spl30_5 | spl30_66), inference(avatar_split_clause, [], [f4070, f678, f387, f433])).
fof(f4070, plain, (~ (e3 = op(e3, e0)) | (~ spl30_5 | spl30_66)), inference(backward_demodulation, [], [f680, f389])).
fof(f4085, plain, (~ spl30_5 | ~ spl30_15 | spl30_149), inference(avatar_contradiction_clause, [], [f4084])).
fof(f4084, plain, ($false | (~ spl30_5 | ~ spl30_15 | spl30_149)), inference(subsumption_resolution, [], [f4076, f431])).
fof(f4076, plain, (~ (e2 = op(e3, e0)) | (~ spl30_5 | spl30_149)), inference(backward_demodulation, [], [f1086, f389])).
fof(f1086, plain, (~ (e2 = op(e3, op(e3, e2))) | spl30_149), inference(avatar_component_clause, [], [f1084])).
fof(f1084, plain, (spl30_149 <=> (e2 = op(e3, op(e3, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl30_149])])).
fof(f4083, plain, (~ spl30_5 | spl30_51 | ~ spl30_148), inference(avatar_contradiction_clause, [], [f4082])).
fof(f4082, plain, ($false | (~ spl30_5 | spl30_51 | ~ spl30_148)), inference(subsumption_resolution, [], [f4075, f583])).
fof(f4075, plain, ((e2 = op(e0, e3)) | (~ spl30_5 | ~ spl30_148)), inference(backward_demodulation, [], [f1081, f389])).
fof(f4080, plain, (~ spl30_5 | ~ spl30_61 | spl30_69), inference(avatar_contradiction_clause, [], [f4079])).
fof(f4079, plain, ($false | (~ spl30_5 | ~ spl30_61 | spl30_69)), inference(subsumption_resolution, [], [f4073, f627])).
fof(f4073, plain, (~ (e0 = op(e0, e0)) | (~ spl30_5 | spl30_69)), inference(backward_demodulation, [], [f695, f389])).
fof(f695, plain, (~ (e0 = op(e0, op(e3, e2))) | spl30_69), inference(avatar_component_clause, [], [f693])).
fof(f4038, plain, (~ spl30_25 | ~ spl30_59 | spl30_165), inference(avatar_contradiction_clause, [], [f4037])).
fof(f4037, plain, ($false | (~ spl30_25 | ~ spl30_59 | spl30_165)), inference(subsumption_resolution, [], [f4032, f618])).
fof(f4032, plain, (~ (e2 = op(e0, e1)) | (~ spl30_25 | spl30_165)), inference(backward_demodulation, [], [f1163, f474])).
fof(f4036, plain, (spl30_16 | ~ spl30_25 | ~ spl30_94), inference(avatar_split_clause, [], [f4029, f814, f472, f433])).
fof(f4029, plain, ((e3 = op(e3, e0)) | (~ spl30_25 | ~ spl30_94)), inference(backward_demodulation, [], [f815, f474])).
fof(f4035, plain, (~ spl30_9 | ~ spl30_25), inference(avatar_split_clause, [], [f4025, f472, f404])).
fof(f4025, plain, (~ (e0 = op(e3, e1)) | ~ spl30_25), inference(backward_demodulation, [], [f174, f474])).
fof(f3961, plain, (~ spl30_53 | ~ spl30_61), inference(avatar_split_clause, [], [f3953, f625, f591])).
fof(f3953, plain, (~ (e0 = op(e0, e2)) | ~ spl30_61), inference(backward_demodulation, [], [f188, f627])).
fof(f3927, plain, (spl30_30 | ~ spl30_59 | ~ spl30_196), inference(avatar_split_clause, [], [f3926, f1310, f616, f493])).
fof(f3926, plain, ((e1 = op(e2, e0)) | (~ spl30_59 | ~ spl30_196)), inference(forward_demodulation, [], [f1312, f618])).
fof(f3921, plain, (spl30_42 | ~ spl30_2 | ~ spl30_202), inference(avatar_split_clause, [], [f3920, f1341, f374, f544])).
fof(f1341, plain, (spl30_202 <=> (e1 = op(e1, op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl30_202])])).
fof(f3920, plain, ((e1 = op(e1, e1)) | (~ spl30_2 | ~ spl30_202)), inference(forward_demodulation, [], [f1342, f376])).
fof(f1342, plain, ((e1 = op(e1, op(e3, e3))) | ~ spl30_202), inference(avatar_component_clause, [], [f1341])).
fof(f3918, plain, (~ spl30_9 | ~ spl30_2 | spl30_223), inference(avatar_split_clause, [], [f3917, f1448, f374, f404])).
fof(f3917, plain, (~ (e0 = op(e3, e1)) | (~ spl30_2 | spl30_223)), inference(forward_demodulation, [], [f1450, f376])).
fof(f3905, plain, (spl30_48 | ~ spl30_33 | ~ spl30_175), inference(avatar_split_clause, [], [f3865, f1209, f506, f569])).
fof(f3865, plain, ((e3 = op(e1, e0)) | (~ spl30_33 | ~ spl30_175)), inference(backward_demodulation, [], [f1210, f508])).
fof(f1210, plain, ((e3 = op(e1, op(e1, e3))) | ~ spl30_175), inference(avatar_component_clause, [], [f1209])).
fof(f3902, plain, (~ spl30_27 | ~ spl30_2 | spl30_201), inference(avatar_split_clause, [], [f3894, f1336, f374, f480])).
fof(f3894, plain, (~ (e2 = op(e2, e1)) | (~ spl30_2 | spl30_201)), inference(backward_demodulation, [], [f1338, f376])).
fof(f3878, plain, (~ spl30_33 | ~ spl30_44 | ~ spl30_47 | spl30_213), inference(avatar_contradiction_clause, [], [f3877])).
fof(f3877, plain, ($false | (~ spl30_33 | ~ spl30_44 | ~ spl30_47 | spl30_213)), inference(subsumption_resolution, [], [f3866, f567])).
fof(f3866, plain, (~ (e2 = op(e1, e0)) | (~ spl30_33 | ~ spl30_44 | spl30_213)), inference(backward_demodulation, [], [f3480, f508])).
fof(f3480, plain, (~ (e2 = op(e1, op(e1, e3))) | (~ spl30_44 | spl30_213)), inference(forward_demodulation, [], [f1399, f554])).
fof(f3876, plain, (~ spl30_33 | ~ spl30_47 | ~ spl30_175), inference(avatar_contradiction_clause, [], [f3875])).
fof(f3875, plain, ($false | (~ spl30_33 | ~ spl30_47 | ~ spl30_175)), inference(subsumption_resolution, [], [f3874, f216])).
fof(f3874, plain, ((e2 = e3) | (~ spl30_33 | ~ spl30_47 | ~ spl30_175)), inference(forward_demodulation, [], [f3865, f567])).
fof(f3849, plain, (spl30_26 | ~ spl30_38 | ~ spl30_98), inference(avatar_split_clause, [], [f3840, f834, f527, f476])).
fof(f3840, plain, ((e1 = op(e2, e1)) | (~ spl30_38 | ~ spl30_98)), inference(backward_demodulation, [], [f836, f529])).
fof(f3832, plain, (~ spl30_37 | ~ spl30_53), inference(avatar_split_clause, [], [f3825, f591, f523])).
fof(f3825, plain, (~ (e0 = op(e1, e2)) | ~ spl30_53), inference(backward_demodulation, [], [f175, f593])).
fof(f3824, plain, (spl30_57 | ~ spl30_62 | ~ spl30_145), inference(avatar_contradiction_clause, [], [f3823])).
fof(f3823, plain, ($false | (spl30_57 | ~ spl30_62 | ~ spl30_145)), inference(subsumption_resolution, [], [f3820, f609])).
fof(f3820, plain, ((e0 = op(e0, e1)) | (~ spl30_62 | ~ spl30_145)), inference(backward_demodulation, [], [f1066, f631])).
fof(f3796, plain, (~ spl30_30 | ~ spl30_47 | spl30_183), inference(avatar_split_clause, [], [f3795, f1248, f565, f493])).
fof(f3795, plain, (~ (e1 = op(e2, e0)) | (~ spl30_47 | spl30_183)), inference(forward_demodulation, [], [f1250, f567])).
fof(f3784, plain, (~ spl30_145 | ~ spl30_208), inference(avatar_contradiction_clause, [], [f3783])).
fof(f3783, plain, ($false | (~ spl30_145 | ~ spl30_208)), inference(subsumption_resolution, [], [f3782, f212])).
fof(f3782, plain, ((e0 = e2) | (~ spl30_145 | ~ spl30_208)), inference(forward_demodulation, [], [f1373, f1066])).
fof(f3750, plain, (~ spl30_38 | ~ spl30_7 | spl30_68), inference(avatar_split_clause, [], [f3715, f688, f395, f527])).
fof(f3715, plain, (~ (e1 = op(e1, e2)) | (~ spl30_7 | spl30_68)), inference(backward_demodulation, [], [f690, f397])).
fof(f397, plain, ((e2 = op(e3, e2)) | ~ spl30_7), inference(avatar_component_clause, [], [f395])).
fof(f3743, plain, (~ spl30_35 | ~ spl30_44 | spl30_214), inference(avatar_split_clause, [], [f3478, f1402, f552, f514])).
fof(f3478, plain, (~ (e2 = op(e1, e3)) | (~ spl30_44 | spl30_214)), inference(forward_demodulation, [], [f1404, f554])).
fof(f3733, plain, (~ spl30_33 | ~ spl30_44 | spl30_211), inference(avatar_split_clause, [], [f3353, f1388, f552, f506])).
fof(f3353, plain, (~ (e0 = op(e1, e3)) | (~ spl30_44 | spl30_211)), inference(backward_demodulation, [], [f1390, f554])).
fof(f1390, plain, (~ (e0 = op(e1, op(e1, e1))) | spl30_211), inference(avatar_component_clause, [], [f1388])).
fof(f3712, plain, (~ spl30_1 | ~ spl30_16 | spl30_157), inference(avatar_contradiction_clause, [], [f3711])).
fof(f3711, plain, ($false | (~ spl30_1 | ~ spl30_16 | spl30_157)), inference(subsumption_resolution, [], [f3707, f372])).
fof(f3707, plain, (~ (e0 = op(e3, e3)) | (~ spl30_16 | spl30_157)), inference(backward_demodulation, [], [f1124, f435])).
fof(f3710, plain, (~ spl30_16 | spl30_155), inference(avatar_contradiction_clause, [], [f3709])).
fof(f3709, plain, ($false | (~ spl30_16 | spl30_155)), inference(subsumption_resolution, [], [f3705, f435])).
fof(f3705, plain, (~ (e3 = op(e3, e0)) | (~ spl30_16 | spl30_155)), inference(backward_demodulation, [], [f1114, f435])).
fof(f3700, plain, (~ spl30_8 | ~ spl30_24), inference(avatar_split_clause, [], [f3696, f467, f399])).
fof(f3696, plain, (~ (e3 = op(e3, e2)) | ~ spl30_24), inference(backward_demodulation, [], [f180, f469])).
fof(f3692, plain, (spl30_25 | ~ spl30_30 | ~ spl30_171), inference(avatar_split_clause, [], [f3691, f1190, f493, f472])).
fof(f3691, plain, ((e0 = op(e2, e1)) | (~ spl30_30 | ~ spl30_171)), inference(backward_demodulation, [], [f1191, f495])).
fof(f3681, plain, (~ spl30_37 | ~ spl30_54 | spl30_177), inference(avatar_contradiction_clause, [], [f3680])).
fof(f3680, plain, ($false | (~ spl30_37 | ~ spl30_54 | spl30_177)), inference(subsumption_resolution, [], [f3672, f597])).
fof(f3672, plain, (~ (e1 = op(e0, e2)) | (~ spl30_37 | spl30_177)), inference(backward_demodulation, [], [f1220, f525])).
fof(f3631, plain, (~ spl30_46 | ~ spl30_1 | spl30_202), inference(avatar_split_clause, [], [f3630, f1341, f370, f561])).
fof(f3630, plain, (~ (e1 = op(e1, e0)) | (~ spl30_1 | spl30_202)), inference(forward_demodulation, [], [f1343, f372])).
fof(f1343, plain, (~ (e1 = op(e1, op(e3, e3))) | spl30_202), inference(avatar_component_clause, [], [f1341])).
fof(f3607, plain, (~ spl30_26 | ~ spl30_10), inference(avatar_split_clause, [], [f3606, f408, f476])).
fof(f3606, plain, (~ (e1 = op(e2, e1)) | ~ spl30_10), inference(forward_demodulation, [], [f174, f410])).
fof(f3604, plain, (~ spl30_27 | ~ spl30_10 | spl30_73), inference(avatar_split_clause, [], [f3555, f712, f408, f480])).
fof(f3555, plain, (~ (e2 = op(e2, e1)) | (~ spl30_10 | spl30_73)), inference(backward_demodulation, [], [f714, f410])).
fof(f3590, plain, (spl30_15 | ~ spl30_1 | ~ spl30_228), inference(avatar_split_clause, [], [f3586, f1472, f370, f429])).
fof(f3586, plain, ((e2 = op(e3, e0)) | (~ spl30_1 | ~ spl30_228)), inference(backward_demodulation, [], [f1473, f372])).
fof(f1473, plain, ((e2 = op(e3, op(e3, e3))) | ~ spl30_228), inference(avatar_component_clause, [], [f1472])).
fof(f3588, plain, (~ spl30_1 | spl30_31 | ~ spl30_201), inference(avatar_contradiction_clause, [], [f3587])).
fof(f3587, plain, ($false | (~ spl30_1 | spl30_31 | ~ spl30_201)), inference(subsumption_resolution, [], [f3582, f498])).
fof(f3582, plain, ((e2 = op(e2, e0)) | (~ spl30_1 | ~ spl30_201)), inference(backward_demodulation, [], [f1337, f372])).
fof(f3576, plain, (~ spl30_4 | ~ spl30_8 | spl30_66), inference(avatar_split_clause, [], [f3568, f678, f399, f382])).
fof(f3568, plain, (~ (e3 = op(e3, e3)) | (~ spl30_8 | spl30_66)), inference(backward_demodulation, [], [f680, f401])).
fof(f401, plain, ((e3 = op(e3, e2)) | ~ spl30_8), inference(avatar_component_clause, [], [f399])).
fof(f3575, plain, (~ spl30_4 | ~ spl30_8), inference(avatar_split_clause, [], [f3567, f399, f382])).
fof(f3567, plain, (~ (e3 = op(e3, e3)) | ~ spl30_8), inference(backward_demodulation, [], [f210, f401])).
fof(f3550, plain, (~ spl30_41 | ~ spl30_44), inference(avatar_contradiction_clause, [], [f3549])).
fof(f3549, plain, ($false | (~ spl30_41 | ~ spl30_44)), inference(subsumption_resolution, [], [f3548, f213])).
fof(f3548, plain, ((e0 = e3) | (~ spl30_41 | ~ spl30_44)), inference(backward_demodulation, [], [f554, f542])).
fof(f3544, plain, (spl30_41 | ~ spl30_46 | ~ spl30_185), inference(avatar_split_clause, [], [f3541, f1258, f561, f540])).
fof(f3541, plain, ((e0 = op(e1, e1)) | (~ spl30_46 | ~ spl30_185)), inference(backward_demodulation, [], [f1259, f563])).
fof(f3538, plain, (~ spl30_52 | spl30_189), inference(avatar_contradiction_clause, [], [f3537])).
fof(f3537, plain, ($false | (~ spl30_52 | spl30_189)), inference(subsumption_resolution, [], [f3531, f588])).
fof(f3531, plain, (~ (e3 = op(e0, e3)) | (~ spl30_52 | spl30_189)), inference(backward_demodulation, [], [f1279, f588])).
fof(f3523, plain, (~ spl30_54 | ~ spl30_59 | spl30_193), inference(avatar_contradiction_clause, [], [f3522])).
fof(f3522, plain, ($false | (~ spl30_54 | ~ spl30_59 | spl30_193)), inference(subsumption_resolution, [], [f3518, f618])).
fof(f3518, plain, (~ (e2 = op(e0, e1)) | (~ spl30_54 | spl30_193)), inference(backward_demodulation, [], [f1298, f597])).
fof(f3511, plain, (~ spl30_61 | ~ spl30_208), inference(avatar_contradiction_clause, [], [f3510])).
fof(f3510, plain, ($false | (~ spl30_61 | ~ spl30_208)), inference(subsumption_resolution, [], [f3509, f212])).
fof(f3509, plain, ((e0 = e2) | (~ spl30_61 | ~ spl30_208)), inference(forward_demodulation, [], [f3501, f627])).
fof(f3501, plain, ((op(e0, e0) = e2) | (~ spl30_61 | ~ spl30_208)), inference(backward_demodulation, [], [f1373, f627])).
fof(f3507, plain, (~ spl30_61 | spl30_145), inference(avatar_contradiction_clause, [], [f3506])).
fof(f3506, plain, ($false | (~ spl30_61 | spl30_145)), inference(subsumption_resolution, [], [f3499, f627])).
fof(f3499, plain, (~ (e0 = op(e0, e0)) | (~ spl30_61 | spl30_145)), inference(backward_demodulation, [], [f1067, f627])).
fof(f3502, plain, (~ spl30_45 | ~ spl30_61), inference(avatar_split_clause, [], [f3490, f625, f557])).
fof(f3490, plain, (~ (e0 = op(e1, e0)) | ~ spl30_61), inference(backward_demodulation, [], [f163, f627])).
fof(f3473, plain, (~ spl30_6 | spl30_225 | ~ spl30_228), inference(avatar_split_clause, [], [f3472, f1472, f1457, f391])).
fof(f3472, plain, (~ (e1 = op(e3, e2)) | (spl30_225 | ~ spl30_228)), inference(forward_demodulation, [], [f1459, f1473])).
fof(f3460, plain, (~ spl30_46 | ~ spl30_34), inference(avatar_split_clause, [], [f3070, f510, f561])).
fof(f3070, plain, (~ (e1 = op(e1, e0)) | ~ spl30_34), inference(forward_demodulation, [], [f195, f512])).
fof(f3454, plain, (~ spl30_12 | ~ spl30_44), inference(avatar_split_clause, [], [f3453, f552, f416])).
fof(f3453, plain, (~ (e3 = op(e3, e1)) | ~ spl30_44), inference(forward_demodulation, [], [f173, f554])).
fof(f3420, plain, (~ spl30_56 | ~ spl30_59 | ~ spl30_62 | spl30_206), inference(avatar_contradiction_clause, [], [f3419])).
fof(f3419, plain, ($false | (~ spl30_56 | ~ spl30_59 | ~ spl30_62 | spl30_206)), inference(subsumption_resolution, [], [f3418, f605])).
fof(f3418, plain, (~ (e3 = op(e0, e2)) | (~ spl30_59 | ~ spl30_62 | spl30_206)), inference(forward_demodulation, [], [f3417, f618])).
fof(f3417, plain, (~ (e3 = op(e0, op(e0, e1))) | (~ spl30_62 | spl30_206)), inference(forward_demodulation, [], [f1364, f631])).
fof(f3410, plain, (~ spl30_21 | ~ spl30_32 | spl30_222), inference(avatar_contradiction_clause, [], [f3409])).
fof(f3409, plain, ($false | (~ spl30_21 | ~ spl30_32 | spl30_222)), inference(subsumption_resolution, [], [f3408, f503])).
fof(f3408, plain, (~ (e3 = op(e2, e0)) | (~ spl30_21 | spl30_222)), inference(forward_demodulation, [], [f1444, f457])).
fof(f1444, plain, (~ (e3 = op(e2, op(e2, e2))) | spl30_222), inference(avatar_component_clause, [], [f1442])).
fof(f3378, plain, (~ spl30_21 | ~ spl30_39 | ~ spl30_98 | ~ spl30_167), inference(avatar_contradiction_clause, [], [f3377])).
fof(f3377, plain, ($false | (~ spl30_21 | ~ spl30_39 | ~ spl30_98 | ~ spl30_167)), inference(subsumption_resolution, [], [f3376, f211])).
fof(f3376, plain, ((e0 = e1) | (~ spl30_21 | ~ spl30_39 | ~ spl30_98 | ~ spl30_167)), inference(forward_demodulation, [], [f3370, f457])).
fof(f3370, plain, ((e1 = op(e2, e2)) | (~ spl30_39 | ~ spl30_98 | ~ spl30_167)), inference(backward_demodulation, [], [f3184, f533])).
fof(f3184, plain, ((e1 = op(e2, op(e1, e2))) | (~ spl30_98 | ~ spl30_167)), inference(forward_demodulation, [], [f1172, f836])).
fof(f3373, plain, (spl30_27 | ~ spl30_39 | ~ spl30_98), inference(avatar_contradiction_clause, [], [f3372])).
fof(f3372, plain, ($false | (spl30_27 | ~ spl30_39 | ~ spl30_98)), inference(subsumption_resolution, [], [f3363, f481])).
fof(f3363, plain, ((e2 = op(e2, e1)) | (~ spl30_39 | ~ spl30_98)), inference(backward_demodulation, [], [f836, f533])).
fof(f3338, plain, (~ spl30_59 | ~ spl30_62 | spl30_208), inference(avatar_split_clause, [], [f3332, f1372, f629, f616])).
fof(f3332, plain, (~ (e2 = op(e0, e1)) | (~ spl30_62 | spl30_208)), inference(backward_demodulation, [], [f1374, f631])).
fof(f3336, plain, (~ spl30_58 | ~ spl30_62), inference(avatar_split_clause, [], [f3326, f629, f612])).
fof(f3326, plain, (~ (e1 = op(e0, e1)) | ~ spl30_62), inference(backward_demodulation, [], [f187, f631])).
fof(f3316, plain, (spl30_64 | ~ spl30_49 | ~ spl30_188), inference(avatar_split_clause, [], [f3225, f1272, f574, f637])).
fof(f3225, plain, ((op(e0, e0) = e3) | (~ spl30_49 | ~ spl30_188)), inference(forward_demodulation, [], [f1274, f576])).
fof(f3312, plain, (~ spl30_60 | ~ spl30_9 | spl30_151), inference(avatar_split_clause, [], [f3287, f1093, f404, f620])).
fof(f3287, plain, (~ (e3 = op(e0, e1)) | (~ spl30_9 | spl30_151)), inference(backward_demodulation, [], [f1095, f406])).
fof(f3296, plain, (~ spl30_31 | ~ spl30_9 | spl30_73), inference(avatar_split_clause, [], [f3284, f712, f404, f497])).
fof(f3284, plain, (~ (e2 = op(e2, e0)) | (~ spl30_9 | spl30_73)), inference(backward_demodulation, [], [f714, f406])).
fof(f3247, plain, (~ spl30_45 | spl30_185), inference(avatar_contradiction_clause, [], [f3246])).
fof(f3246, plain, ($false | (~ spl30_45 | spl30_185)), inference(subsumption_resolution, [], [f3242, f559])).
fof(f3242, plain, (~ (e0 = op(e1, e0)) | (~ spl30_45 | spl30_185)), inference(backward_demodulation, [], [f1260, f559])).
fof(f3245, plain, (~ spl30_45 | spl30_57 | ~ spl30_184), inference(avatar_contradiction_clause, [], [f3244])).
fof(f3244, plain, ($false | (~ spl30_45 | spl30_57 | ~ spl30_184)), inference(subsumption_resolution, [], [f3241, f609])).
fof(f3241, plain, ((e0 = op(e0, e1)) | (~ spl30_45 | ~ spl30_184)), inference(backward_demodulation, [], [f1255, f559])).
fof(f3243, plain, (~ spl30_37 | ~ spl30_45), inference(avatar_split_clause, [], [f3236, f557, f523])).
fof(f3236, plain, (~ (e0 = op(e1, e2)) | ~ spl30_45), inference(backward_demodulation, [], [f194, f559])).
fof(f3216, plain, (~ spl30_25 | ~ spl30_220 | spl30_221), inference(avatar_split_clause, [], [f3213, f1437, f1432, f472])).
fof(f3213, plain, (~ (e0 = op(e2, e1)) | (~ spl30_220 | spl30_221)), inference(backward_demodulation, [], [f1439, f1433])).
fof(f3209, plain, (~ spl30_4 | spl30_200), inference(avatar_contradiction_clause, [], [f3208])).
fof(f3208, plain, ($false | (~ spl30_4 | spl30_200)), inference(subsumption_resolution, [], [f3207, f384])).
fof(f3207, plain, (~ (e3 = op(e3, e3)) | (~ spl30_4 | spl30_200)), inference(forward_demodulation, [], [f1333, f384])).
fof(f1333, plain, (~ (e3 = op(e3, op(e3, e3))) | spl30_200), inference(avatar_component_clause, [], [f1331])).
fof(f3197, plain, (~ spl30_37 | ~ spl30_43 | spl30_211), inference(avatar_split_clause, [], [f3083, f1388, f548, f523])).
fof(f3083, plain, (~ (e0 = op(e1, e2)) | (~ spl30_43 | spl30_211)), inference(forward_demodulation, [], [f1390, f550])).
fof(f3186, plain, (~ spl30_27 | ~ spl30_6 | spl30_67), inference(avatar_split_clause, [], [f3039, f683, f391, f480])).
fof(f683, plain, (spl30_67 <=> (e2 = op(e2, op(e3, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl30_67])])).
fof(f3039, plain, (~ (e2 = op(e2, e1)) | (~ spl30_6 | spl30_67)), inference(forward_demodulation, [], [f685, f393])).
fof(f685, plain, (~ (e2 = op(e2, op(e3, e2))) | spl30_67), inference(avatar_component_clause, [], [f683])).
fof(f3178, plain, (spl30_24 | ~ spl30_19 | ~ spl30_160), inference(avatar_split_clause, [], [f3095, f1136, f446, f467])).
fof(f3095, plain, ((e3 = op(e2, e2)) | (~ spl30_19 | ~ spl30_160)), inference(forward_demodulation, [], [f1138, f448])).
fof(f3172, plain, (~ spl30_4 | ~ spl30_49 | spl30_203), inference(avatar_contradiction_clause, [], [f3171])).
fof(f3171, plain, ($false | (~ spl30_4 | ~ spl30_49 | spl30_203)), inference(subsumption_resolution, [], [f3164, f576])).
fof(f3164, plain, (~ (e0 = op(e0, e3)) | (~ spl30_4 | spl30_203)), inference(backward_demodulation, [], [f1348, f384])).
fof(f3168, plain, (~ spl30_4 | ~ spl30_19 | spl30_201), inference(avatar_contradiction_clause, [], [f3167])).
fof(f3167, plain, ($false | (~ spl30_4 | ~ spl30_19 | spl30_201)), inference(subsumption_resolution, [], [f3162, f448])).
fof(f3162, plain, (~ (e2 = op(e2, e3)) | (~ spl30_4 | spl30_201)), inference(backward_demodulation, [], [f1338, f384])).
fof(f3146, plain, (~ spl30_13 | ~ spl30_15), inference(avatar_contradiction_clause, [], [f3145])).
fof(f3145, plain, ($false | (~ spl30_13 | ~ spl30_15)), inference(subsumption_resolution, [], [f3144, f212])).
fof(f3144, plain, ((e0 = e2) | (~ spl30_13 | ~ spl30_15)), inference(backward_demodulation, [], [f431, f423])).
fof(f3142, plain, (~ spl30_43 | ~ spl30_44), inference(avatar_contradiction_clause, [], [f3141])).
fof(f3141, plain, ($false | (~ spl30_43 | ~ spl30_44)), inference(subsumption_resolution, [], [f3140, f216])).
fof(f3140, plain, ((e2 = e3) | (~ spl30_43 | ~ spl30_44)), inference(forward_demodulation, [], [f554, f550])).
fof(f3128, plain, (~ spl30_1 | ~ spl30_49), inference(avatar_split_clause, [], [f3116, f574, f370])).
fof(f3116, plain, (~ (e0 = op(e3, e3)) | ~ spl30_49), inference(backward_demodulation, [], [f183, f576])).
fof(f3101, plain, (~ spl30_61 | ~ spl30_37 | spl30_114), inference(avatar_split_clause, [], [f3100, f912, f523, f625])).
fof(f3100, plain, (~ (e0 = op(e0, e0)) | (~ spl30_37 | spl30_114)), inference(forward_demodulation, [], [f914, f525])).
fof(f3098, plain, (~ spl30_19 | spl30_159), inference(avatar_contradiction_clause, [], [f3097])).
fof(f3097, plain, ($false | (~ spl30_19 | spl30_159)), inference(subsumption_resolution, [], [f3096, f448])).
fof(f3096, plain, (~ (e2 = op(e2, e3)) | (~ spl30_19 | spl30_159)), inference(forward_demodulation, [], [f1133, f448])).
fof(f1133, plain, (~ (e2 = op(op(e2, e3), e3)) | spl30_159), inference(avatar_component_clause, [], [f1131])).
fof(f3053, plain, (spl30_12 | ~ spl30_6 | ~ spl30_66), inference(avatar_split_clause, [], [f3041, f678, f391, f416])).
fof(f3041, plain, ((e3 = op(e3, e1)) | (~ spl30_6 | ~ spl30_66)), inference(forward_demodulation, [], [f679, f393])).
fof(f679, plain, ((e3 = op(e3, op(e3, e2))) | ~ spl30_66), inference(avatar_component_clause, [], [f678])).
fof(f3038, plain, (~ spl30_6 | spl30_18 | ~ spl30_70), inference(avatar_contradiction_clause, [], [f3037])).
fof(f3037, plain, ($false | (~ spl30_6 | spl30_18 | ~ spl30_70)), inference(subsumption_resolution, [], [f3036, f443])).
fof(f3036, plain, ((e1 = op(e2, e3)) | (~ spl30_6 | ~ spl30_70)), inference(forward_demodulation, [], [f700, f393])).
fof(f3021, plain, (~ spl30_25 | spl30_54 | ~ spl30_166), inference(avatar_contradiction_clause, [], [f3020])).
fof(f3020, plain, ($false | (~ spl30_25 | spl30_54 | ~ spl30_166)), inference(subsumption_resolution, [], [f3019, f596])).
fof(f3019, plain, ((e1 = op(e0, e2)) | (~ spl30_25 | ~ spl30_166)), inference(forward_demodulation, [], [f1168, f474])).
fof(f3001, plain, (~ spl30_1 | ~ spl30_6 | ~ spl30_15 | spl30_225), inference(avatar_contradiction_clause, [], [f3000])).
fof(f3000, plain, ($false | (~ spl30_1 | ~ spl30_6 | ~ spl30_15 | spl30_225)), inference(subsumption_resolution, [], [f2999, f393])).
fof(f2999, plain, (~ (e1 = op(e3, e2)) | (~ spl30_1 | ~ spl30_15 | spl30_225)), inference(forward_demodulation, [], [f2998, f431])).
fof(f2998, plain, (~ (e1 = op(e3, op(e3, e0))) | (~ spl30_1 | spl30_225)), inference(forward_demodulation, [], [f1459, f372])).
fof(f2980, plain, (~ spl30_25 | ~ spl30_61 | spl30_97), inference(avatar_contradiction_clause, [], [f2979])).
fof(f2979, plain, ($false | (~ spl30_25 | ~ spl30_61 | spl30_97)), inference(subsumption_resolution, [], [f2978, f627])).
fof(f2978, plain, (~ (e0 = op(e0, e0)) | (~ spl30_25 | spl30_97)), inference(forward_demodulation, [], [f831, f474])).
fof(f831, plain, (~ (e0 = op(e0, op(e2, e1))) | spl30_97), inference(avatar_component_clause, [], [f829])).
fof(f2962, plain, (~ spl30_12 | spl30_49 | ~ spl30_75), inference(avatar_contradiction_clause, [], [f2961])).
fof(f2961, plain, ($false | (~ spl30_12 | spl30_49 | ~ spl30_75)), inference(subsumption_resolution, [], [f2960, f575])).
fof(f2960, plain, ((e0 = op(e0, e3)) | (~ spl30_12 | ~ spl30_75)), inference(forward_demodulation, [], [f723, f418])).
fof(f723, plain, ((e0 = op(e0, op(e3, e1))) | ~ spl30_75), inference(avatar_component_clause, [], [f722])).
fof(f2959, plain, (~ spl30_12 | spl30_151), inference(avatar_contradiction_clause, [], [f2958])).
fof(f2958, plain, ($false | (~ spl30_12 | spl30_151)), inference(subsumption_resolution, [], [f2957, f418])).
fof(f2957, plain, (~ (e3 = op(e3, e1)) | (~ spl30_12 | spl30_151)), inference(forward_demodulation, [], [f1095, f418])).
fof(f2949, plain, (spl30_47 | ~ spl30_37 | ~ spl30_179), inference(avatar_split_clause, [], [f2865, f1228, f523, f565])).
fof(f2865, plain, ((e2 = op(e1, e0)) | (~ spl30_37 | ~ spl30_179)), inference(backward_demodulation, [], [f1229, f525])).
fof(f2948, plain, (~ spl30_1 | ~ spl30_15 | spl30_228), inference(avatar_contradiction_clause, [], [f2947])).
fof(f2947, plain, ($false | (~ spl30_1 | ~ spl30_15 | spl30_228)), inference(subsumption_resolution, [], [f2938, f431])).
fof(f2938, plain, (~ (e2 = op(e3, e0)) | (~ spl30_1 | spl30_228)), inference(backward_demodulation, [], [f1474, f372])).
fof(f2943, plain, (~ spl30_1 | ~ spl30_61 | spl30_203), inference(avatar_contradiction_clause, [], [f2942])).
fof(f2942, plain, ($false | (~ spl30_1 | ~ spl30_61 | spl30_203)), inference(subsumption_resolution, [], [f2935, f627])).
fof(f2935, plain, (~ (e0 = op(e0, e0)) | (~ spl30_1 | spl30_203)), inference(backward_demodulation, [], [f1348, f372])).
fof(f2920, plain, (~ spl30_2 | ~ spl30_6), inference(avatar_split_clause, [], [f2911, f391, f374])).
fof(f2911, plain, (~ (e1 = op(e3, e3)) | ~ spl30_6), inference(backward_demodulation, [], [f210, f393])).
fof(f2888, plain, (~ spl30_34 | spl30_173), inference(avatar_contradiction_clause, [], [f2887])).
fof(f2887, plain, ($false | (~ spl30_34 | spl30_173)), inference(subsumption_resolution, [], [f2879, f512])).
fof(f2879, plain, (~ (e1 = op(e1, e3)) | (~ spl30_34 | spl30_173)), inference(backward_demodulation, [], [f1201, f512])).
fof(f2801, plain, (spl30_37 | ~ spl30_30 | ~ spl30_170), inference(avatar_split_clause, [], [f2761, f1185, f493, f523])).
fof(f2761, plain, ((e0 = op(e1, e2)) | (~ spl30_30 | ~ spl30_170)), inference(backward_demodulation, [], [f1187, f495])).
fof(f2797, plain, (spl30_33 | ~ spl30_48 | ~ spl30_185), inference(avatar_split_clause, [], [f2043, f1258, f569, f506])).
fof(f2043, plain, ((e0 = op(e1, e3)) | (~ spl30_48 | ~ spl30_185)), inference(forward_demodulation, [], [f1259, f571])).
fof(f2776, plain, (~ spl30_11 | ~ spl30_15), inference(avatar_split_clause, [], [f2775, f429, f412])).
fof(f2775, plain, (~ (e2 = op(e3, e1)) | ~ spl30_15), inference(forward_demodulation, [], [f205, f431])).
fof(f2763, plain, (~ spl30_22 | ~ spl30_30), inference(avatar_split_clause, [], [f2757, f493, f459])).
fof(f2757, plain, (~ (e1 = op(e2, e2)) | ~ spl30_30), inference(backward_demodulation, [], [f200, f495])).
fof(f2756, plain, (spl30_37 | ~ spl30_54 | ~ spl30_191), inference(avatar_contradiction_clause, [], [f2755])).
fof(f2755, plain, ($false | (spl30_37 | ~ spl30_54 | ~ spl30_191)), inference(subsumption_resolution, [], [f2753, f524])).
fof(f524, plain, (~ (e0 = op(e1, e2)) | spl30_37), inference(avatar_component_clause, [], [f523])).
fof(f2753, plain, ((e0 = op(e1, e2)) | (~ spl30_54 | ~ spl30_191)), inference(backward_demodulation, [], [f1287, f597])).
fof(f2730, plain, (~ spl30_15 | spl30_51 | ~ spl30_82), inference(avatar_contradiction_clause, [], [f2729])).
fof(f2729, plain, ($false | (~ spl30_15 | spl30_51 | ~ spl30_82)), inference(subsumption_resolution, [], [f2728, f583])).
fof(f2728, plain, ((e2 = op(e0, e3)) | (~ spl30_15 | ~ spl30_82)), inference(forward_demodulation, [], [f758, f431])).
fof(f2694, plain, (~ spl30_2 | ~ spl30_9 | ~ spl30_15 | spl30_224), inference(avatar_contradiction_clause, [], [f2693])).
fof(f2693, plain, ($false | (~ spl30_2 | ~ spl30_9 | ~ spl30_15 | spl30_224)), inference(subsumption_resolution, [], [f2692, f431])).
fof(f2692, plain, (~ (e2 = op(e3, e0)) | (~ spl30_2 | ~ spl30_9 | spl30_224)), inference(forward_demodulation, [], [f2691, f406])).
fof(f2691, plain, (~ (e2 = op(e3, op(e3, e1))) | (~ spl30_2 | spl30_224)), inference(forward_demodulation, [], [f1454, f376])).
fof(f2676, plain, (spl30_22 | ~ spl30_39 | ~ spl30_177), inference(avatar_split_clause, [], [f2552, f1218, f531, f459])).
fof(f2552, plain, ((e1 = op(e2, e2)) | (~ spl30_39 | ~ spl30_177)), inference(forward_demodulation, [], [f1219, f533])).
fof(f2673, plain, (~ spl30_2 | spl30_57 | ~ spl30_203), inference(avatar_contradiction_clause, [], [f2672])).
fof(f2672, plain, ($false | (~ spl30_2 | spl30_57 | ~ spl30_203)), inference(subsumption_resolution, [], [f2670, f609])).
fof(f2670, plain, ((e0 = op(e0, e1)) | (~ spl30_2 | ~ spl30_203)), inference(backward_demodulation, [], [f1347, f376])).
fof(f2663, plain, (~ spl30_15 | spl30_17 | ~ spl30_156), inference(avatar_contradiction_clause, [], [f2662])).
fof(f2662, plain, ($false | (~ spl30_15 | spl30_17 | ~ spl30_156)), inference(subsumption_resolution, [], [f2656, f439])).
fof(f2656, plain, ((e0 = op(e2, e3)) | (~ spl30_15 | ~ spl30_156)), inference(backward_demodulation, [], [f1119, f431])).
fof(f2639, plain, (spl30_6 | ~ spl30_28 | ~ spl30_166), inference(avatar_contradiction_clause, [], [f2638])).
fof(f2638, plain, ($false | (spl30_6 | ~ spl30_28 | ~ spl30_166)), inference(subsumption_resolution, [], [f2632, f392])).
fof(f392, plain, (~ (e1 = op(e3, e2)) | spl30_6), inference(avatar_component_clause, [], [f391])).
fof(f2632, plain, ((e1 = op(e3, e2)) | (~ spl30_28 | ~ spl30_166)), inference(backward_demodulation, [], [f1168, f486])).
fof(f2636, plain, (~ spl30_20 | ~ spl30_28), inference(avatar_split_clause, [], [f2630, f484, f450])).
fof(f2630, plain, (~ (e3 = op(e2, e3)) | ~ spl30_28), inference(backward_demodulation, [], [f203, f486])).
fof(f2617, plain, (~ spl30_29 | ~ spl30_59 | ~ spl30_196), inference(avatar_contradiction_clause, [], [f2616])).
fof(f2616, plain, ($false | (~ spl30_29 | ~ spl30_59 | ~ spl30_196)), inference(subsumption_resolution, [], [f2615, f211])).
fof(f2615, plain, ((e0 = e1) | (~ spl30_29 | ~ spl30_59 | ~ spl30_196)), inference(forward_demodulation, [], [f2610, f491])).
fof(f2610, plain, ((e1 = op(e2, e0)) | (~ spl30_59 | ~ spl30_196)), inference(backward_demodulation, [], [f1312, f618])).
fof(f2605, plain, (~ spl30_60 | ~ spl30_62 | spl30_210), inference(avatar_split_clause, [], [f2602, f1382, f629, f620])).
fof(f2602, plain, (~ (e3 = op(e0, e1)) | (~ spl30_62 | spl30_210)), inference(backward_demodulation, [], [f1384, f631])).
fof(f2592, plain, (~ spl30_3 | ~ spl30_8 | spl30_148), inference(avatar_split_clause, [], [f2591, f1079, f399, f378])).
fof(f2591, plain, (~ (e2 = op(e3, e3)) | (~ spl30_8 | spl30_148)), inference(forward_demodulation, [], [f1080, f401])).
fof(f1080, plain, (~ (e2 = op(op(e3, e2), e3)) | spl30_148), inference(avatar_component_clause, [], [f1079])).
fof(f2590, plain, (~ spl30_3 | ~ spl30_8 | spl30_149), inference(avatar_split_clause, [], [f2589, f1084, f399, f378])).
fof(f2589, plain, (~ (e2 = op(e3, e3)) | (~ spl30_8 | spl30_149)), inference(forward_demodulation, [], [f1086, f401])).
fof(f2573, plain, (spl30_63 | ~ spl30_29 | ~ spl30_169), inference(avatar_split_clause, [], [f2516, f1180, f489, f633])).
fof(f2516, plain, ((op(e0, e0) = e2) | (~ spl30_29 | ~ spl30_169)), inference(forward_demodulation, [], [f1181, f491])).
fof(f1181, plain, ((e2 = op(op(e2, e0), e0)) | ~ spl30_169), inference(avatar_component_clause, [], [f1180])).
fof(f2572, plain, (spl30_63 | ~ spl30_53 | ~ spl30_192), inference(avatar_split_clause, [], [f2547, f1291, f591, f633])).
fof(f2547, plain, ((op(e0, e0) = e2) | (~ spl30_53 | ~ spl30_192)), inference(forward_demodulation, [], [f1293, f593])).
fof(f2570, plain, (spl30_60 | ~ spl30_33 | ~ spl30_174), inference(avatar_split_clause, [], [f2522, f1204, f506, f620])).
fof(f2522, plain, ((e3 = op(e0, e1)) | (~ spl30_33 | ~ spl30_174)), inference(forward_demodulation, [], [f1206, f508])).
fof(f2569, plain, (spl30_50 | ~ spl30_33 | ~ spl30_173), inference(avatar_split_clause, [], [f2527, f1199, f506, f578])).
fof(f2527, plain, ((e1 = op(e0, e3)) | (~ spl30_33 | ~ spl30_173)), inference(forward_demodulation, [], [f1200, f508])).
fof(f1200, plain, ((e1 = op(op(e1, e3), e3)) | ~ spl30_173), inference(avatar_component_clause, [], [f1199])).
fof(f2563, plain, (spl30_14 | ~ spl30_9 | ~ spl30_153), inference(avatar_split_clause, [], [f2318, f1103, f404, f425])).
fof(f2318, plain, ((e1 = op(e3, e0)) | (~ spl30_9 | ~ spl30_153)), inference(forward_demodulation, [], [f1104, f406])).
fof(f2558, plain, (~ spl30_1 | ~ spl30_9), inference(avatar_split_clause, [], [f2557, f404, f370])).
fof(f2557, plain, (~ (e0 = op(e3, e3)) | ~ spl30_9), inference(forward_demodulation, [], [f209, f406])).
fof(f2553, plain, (~ spl30_1 | ~ spl30_33), inference(avatar_split_clause, [], [f1984, f506, f370])).
fof(f1984, plain, (~ (e0 = op(e3, e3)) | ~ spl30_33), inference(forward_demodulation, [], [f185, f508])).
fof(f2550, plain, (~ spl30_53 | spl30_191), inference(avatar_contradiction_clause, [], [f2549])).
fof(f2549, plain, ($false | (~ spl30_53 | spl30_191)), inference(subsumption_resolution, [], [f2548, f593])).
fof(f2548, plain, (~ (e0 = op(e0, e2)) | (~ spl30_53 | spl30_191)), inference(forward_demodulation, [], [f1288, f593])).
fof(f2537, plain, (~ spl30_39 | spl30_179), inference(avatar_contradiction_clause, [], [f2536])).
fof(f2536, plain, ($false | (~ spl30_39 | spl30_179)), inference(subsumption_resolution, [], [f2535, f533])).
fof(f2535, plain, (~ (e2 = op(e1, e2)) | (~ spl30_39 | spl30_179)), inference(forward_demodulation, [], [f1230, f533])).
fof(f2534, plain, (~ spl30_27 | spl30_165), inference(avatar_contradiction_clause, [], [f2533])).
fof(f2533, plain, ($false | (~ spl30_27 | spl30_165)), inference(subsumption_resolution, [], [f2532, f482])).
fof(f2532, plain, (~ (e2 = op(e2, e1)) | (~ spl30_27 | spl30_165)), inference(forward_demodulation, [], [f1163, f482])).
fof(f2525, plain, (~ spl30_33 | ~ spl30_50 | spl30_173), inference(avatar_contradiction_clause, [], [f2524])).
fof(f2524, plain, ($false | (~ spl30_33 | ~ spl30_50 | spl30_173)), inference(subsumption_resolution, [], [f2523, f580])).
fof(f2523, plain, (~ (e1 = op(e0, e3)) | (~ spl30_33 | spl30_173)), inference(forward_demodulation, [], [f1201, f508])).
fof(f2521, plain, (~ spl30_33 | ~ spl30_48 | spl30_175), inference(avatar_contradiction_clause, [], [f2520])).
fof(f2520, plain, ($false | (~ spl30_33 | ~ spl30_48 | spl30_175)), inference(subsumption_resolution, [], [f2519, f571])).
fof(f2519, plain, (~ (e3 = op(e1, e0)) | (~ spl30_33 | spl30_175)), inference(forward_demodulation, [], [f1211, f508])).
fof(f2503, plain, (~ spl30_29 | spl30_171), inference(avatar_contradiction_clause, [], [f2502])).
fof(f2502, plain, ($false | (~ spl30_29 | spl30_171)), inference(subsumption_resolution, [], [f2501, f491])).
fof(f2501, plain, (~ (e0 = op(e2, e0)) | (~ spl30_29 | spl30_171)), inference(forward_demodulation, [], [f1192, f491])).
fof(f2485, plain, (~ spl30_20 | spl30_161), inference(avatar_contradiction_clause, [], [f2484])).
fof(f2484, plain, ($false | (~ spl30_20 | spl30_161)), inference(subsumption_resolution, [], [f2481, f452])).
fof(f2481, plain, (~ (e3 = op(e2, e3)) | (~ spl30_20 | spl30_161)), inference(backward_demodulation, [], [f1143, f452])).
fof(f2483, plain, (~ spl30_3 | ~ spl30_20 | spl30_159), inference(avatar_contradiction_clause, [], [f2482])).
fof(f2482, plain, ($false | (~ spl30_3 | ~ spl30_20 | spl30_159)), inference(subsumption_resolution, [], [f2479, f380])).
fof(f2479, plain, (~ (e2 = op(e3, e3)) | (~ spl30_20 | spl30_159)), inference(backward_demodulation, [], [f1133, f452])).
fof(f2457, plain, (~ spl30_21 | ~ spl30_29), inference(avatar_split_clause, [], [f2453, f489, f455])).
fof(f2453, plain, (~ (e0 = op(e2, e2)) | ~ spl30_29), inference(backward_demodulation, [], [f200, f491])).
fof(f2445, plain, (~ spl30_21 | ~ spl30_53), inference(avatar_split_clause, [], [f2438, f591, f455])).
fof(f2438, plain, (~ (e0 = op(e2, e2)) | ~ spl30_53), inference(backward_demodulation, [], [f176, f593])).
fof(f2432, plain, (~ spl30_54 | ~ spl30_63 | spl30_205), inference(avatar_split_clause, [], [f2426, f1358, f633, f595])).
fof(f2426, plain, (~ (e1 = op(e0, e2)) | (~ spl30_63 | spl30_205)), inference(backward_demodulation, [], [f1360, f635])).
fof(f1360, plain, (~ (e1 = op(e0, op(e0, e0))) | spl30_205), inference(avatar_component_clause, [], [f1358])).
fof(f2420, plain, (spl30_60 | ~ spl30_9 | ~ spl30_151), inference(avatar_split_clause, [], [f2419, f1093, f404, f620])).
fof(f2419, plain, ((e3 = op(e0, e1)) | (~ spl30_9 | ~ spl30_151)), inference(forward_demodulation, [], [f1094, f406])).
fof(f2372, plain, (~ spl30_61 | ~ spl30_9 | spl30_75), inference(avatar_split_clause, [], [f2371, f722, f404, f625])).
fof(f2371, plain, (~ (e0 = op(e0, e0)) | (~ spl30_9 | spl30_75)), inference(forward_demodulation, [], [f724, f406])).
fof(f2355, plain, (~ spl30_18 | spl30_219 | ~ spl30_222), inference(avatar_contradiction_clause, [], [f2354])).
fof(f2354, plain, ($false | (~ spl30_18 | spl30_219 | ~ spl30_222)), inference(subsumption_resolution, [], [f2353, f444])).
fof(f2353, plain, (~ (e1 = op(e2, e3)) | (spl30_219 | ~ spl30_222)), inference(forward_demodulation, [], [f1429, f1443])).
fof(f2327, plain, (~ spl30_27 | ~ spl30_28), inference(avatar_contradiction_clause, [], [f2326])).
fof(f2326, plain, ($false | (~ spl30_27 | ~ spl30_28)), inference(subsumption_resolution, [], [f2325, f216])).
fof(f2325, plain, ((e2 = e3) | (~ spl30_27 | ~ spl30_28)), inference(forward_demodulation, [], [f486, f482])).
fof(f2322, plain, (~ spl30_50 | ~ spl30_58), inference(avatar_split_clause, [], [f2320, f612, f578])).
fof(f2320, plain, (~ (e1 = op(e0, e3)) | ~ spl30_58), inference(backward_demodulation, [], [f191, f614])).
fof(f2297, plain, (spl30_50 | ~ spl30_9 | ~ spl30_152), inference(avatar_split_clause, [], [f2286, f1098, f404, f578])).
fof(f2286, plain, ((e1 = op(e0, e3)) | (~ spl30_9 | ~ spl30_152)), inference(backward_demodulation, [], [f1100, f406])).
fof(f2292, plain, (spl30_20 | ~ spl30_8 | ~ spl30_70), inference(avatar_split_clause, [], [f2149, f698, f399, f450])).
fof(f2149, plain, ((e3 = op(e2, e3)) | (~ spl30_8 | ~ spl30_70)), inference(forward_demodulation, [], [f700, f401])).
fof(f2289, plain, (~ spl30_9 | ~ spl30_14 | spl30_153), inference(avatar_contradiction_clause, [], [f2288])).
fof(f2288, plain, ($false | (~ spl30_9 | ~ spl30_14 | spl30_153)), inference(subsumption_resolution, [], [f2285, f427])).
fof(f2285, plain, (~ (e1 = op(e3, e0)) | (~ spl30_9 | spl30_153)), inference(backward_demodulation, [], [f1105, f406])).
fof(f2272, plain, (~ spl30_25 | ~ spl30_22 | spl30_217), inference(avatar_split_clause, [], [f2271, f1418, f459, f472])).
fof(f2271, plain, (~ (e0 = op(e2, e1)) | (~ spl30_22 | spl30_217)), inference(forward_demodulation, [], [f1420, f461])).
fof(f2270, plain, (~ spl30_31 | ~ spl30_27), inference(avatar_split_clause, [], [f2269, f480, f497])).
fof(f2269, plain, (~ (e2 = op(e2, e0)) | ~ spl30_27), inference(forward_demodulation, [], [f199, f482])).
fof(f2268, plain, (~ spl30_15 | ~ spl30_63), inference(avatar_split_clause, [], [f2267, f633, f429])).
fof(f2267, plain, (~ (e2 = op(e3, e0)) | ~ spl30_63), inference(forward_demodulation, [], [f165, f635])).
fof(f2258, plain, (~ spl30_10 | spl30_153), inference(avatar_contradiction_clause, [], [f2257])).
fof(f2257, plain, ($false | (~ spl30_10 | spl30_153)), inference(subsumption_resolution, [], [f2254, f410])).
fof(f2254, plain, (~ (e1 = op(e3, e1)) | (~ spl30_10 | spl30_153)), inference(backward_demodulation, [], [f1105, f410])).
fof(f2251, plain, (~ spl30_25 | ~ spl30_27), inference(avatar_contradiction_clause, [], [f2250])).
fof(f2250, plain, ($false | (~ spl30_25 | ~ spl30_27)), inference(subsumption_resolution, [], [f2249, f212])).
fof(f2249, plain, ((e0 = e2) | (~ spl30_25 | ~ spl30_27)), inference(backward_demodulation, [], [f482, f474])).
fof(f2237, plain, (~ spl30_49 | ~ spl30_53), inference(avatar_split_clause, [], [f2234, f591, f574])).
fof(f2234, plain, (~ (e0 = op(e0, e3)) | ~ spl30_53), inference(backward_demodulation, [], [f192, f593])).
fof(f2207, plain, (~ spl30_2 | ~ spl30_3), inference(avatar_contradiction_clause, [], [f2206])).
fof(f2206, plain, ($false | (~ spl30_2 | ~ spl30_3)), inference(subsumption_resolution, [], [f2205, f214])).
fof(f2205, plain, ((e1 = e2) | (~ spl30_2 | ~ spl30_3)), inference(backward_demodulation, [], [f380, f376])).
fof(f2192, plain, (~ spl30_21 | ~ spl30_22), inference(avatar_contradiction_clause, [], [f2191])).
fof(f2191, plain, ($false | (~ spl30_21 | ~ spl30_22)), inference(subsumption_resolution, [], [f2190, f211])).
fof(f2190, plain, ((e0 = e1) | (~ spl30_21 | ~ spl30_22)), inference(backward_demodulation, [], [f461, f457])).
fof(f2181, plain, (~ spl30_46 | ~ spl30_48), inference(avatar_contradiction_clause, [], [f2180])).
fof(f2180, plain, ($false | (~ spl30_46 | ~ spl30_48)), inference(subsumption_resolution, [], [f2179, f215])).
fof(f2179, plain, ((e1 = e3) | (~ spl30_46 | ~ spl30_48)), inference(backward_demodulation, [], [f571, f563])).
fof(f2169, plain, (~ spl30_59 | ~ spl30_60), inference(avatar_contradiction_clause, [], [f2168])).
fof(f2168, plain, ($false | (~ spl30_59 | ~ spl30_60)), inference(subsumption_resolution, [], [f2167, f216])).
fof(f2167, plain, ((e2 = e3) | (~ spl30_59 | ~ spl30_60)), inference(backward_demodulation, [], [f622, f618])).
fof(f2165, plain, (~ spl30_50 | ~ spl30_64 | spl30_205), inference(avatar_split_clause, [], [f2160, f1358, f637, f578])).
fof(f2160, plain, (~ (e1 = op(e0, e3)) | (~ spl30_64 | spl30_205)), inference(backward_demodulation, [], [f1360, f639])).
fof(f2151, plain, (~ spl30_49 | ~ spl30_8 | spl30_69), inference(avatar_split_clause, [], [f2150, f693, f399, f574])).
fof(f2150, plain, (~ (e0 = op(e0, e3)) | (~ spl30_8 | spl30_69)), inference(forward_demodulation, [], [f695, f401])).
fof(f2130, plain, (~ spl30_64 | ~ spl30_48), inference(avatar_split_clause, [], [f2036, f569, f637])).
fof(f2036, plain, (~ (op(e0, e0) = e3) | ~ spl30_48), inference(forward_demodulation, [], [f163, f571])).
fof(f2112, plain, (~ spl30_28 | ~ spl30_22 | spl30_222), inference(avatar_split_clause, [], [f2111, f1442, f459, f484])).
fof(f2111, plain, (~ (e3 = op(e2, e1)) | (~ spl30_22 | spl30_222)), inference(forward_demodulation, [], [f1444, f461])).
fof(f2093, plain, (~ spl30_9 | ~ spl30_13), inference(avatar_split_clause, [], [f2092, f421, f404])).
fof(f2092, plain, (~ (e0 = op(e3, e1)) | ~ spl30_13), inference(forward_demodulation, [], [f205, f423])).
fof(f2075, plain, (spl30_3 | ~ spl30_8 | ~ spl30_149), inference(avatar_split_clause, [], [f2072, f1084, f399, f378])).
fof(f2072, plain, ((e2 = op(e3, e3)) | (~ spl30_8 | ~ spl30_149)), inference(backward_demodulation, [], [f1085, f401])).
fof(f1085, plain, ((e2 = op(e3, op(e3, e2))) | ~ spl30_149), inference(avatar_component_clause, [], [f1084])).
fof(f2074, plain, (spl30_3 | ~ spl30_8 | ~ spl30_148), inference(avatar_split_clause, [], [f2071, f1079, f399, f378])).
fof(f2071, plain, ((e2 = op(e3, e3)) | (~ spl30_8 | ~ spl30_148)), inference(backward_demodulation, [], [f1081, f401])).
fof(f2068, plain, (~ spl30_13 | ~ spl30_14), inference(avatar_contradiction_clause, [], [f2067])).
fof(f2067, plain, ($false | (~ spl30_13 | ~ spl30_14)), inference(subsumption_resolution, [], [f2066, f211])).
fof(f2066, plain, ((e0 = e1) | (~ spl30_13 | ~ spl30_14)), inference(forward_demodulation, [], [f427, f423])).
fof(f2065, plain, (~ spl30_18 | ~ spl30_20), inference(avatar_contradiction_clause, [], [f2064])).
fof(f2064, plain, ($false | (~ spl30_18 | ~ spl30_20)), inference(subsumption_resolution, [], [f2063, f215])).
fof(f2063, plain, ((e1 = e3) | (~ spl30_18 | ~ spl30_20)), inference(forward_demodulation, [], [f452, f444])).
fof(f2062, plain, (~ spl30_22 | ~ spl30_23), inference(avatar_contradiction_clause, [], [f2061])).
fof(f2061, plain, ($false | (~ spl30_22 | ~ spl30_23)), inference(subsumption_resolution, [], [f2060, f214])).
fof(f2060, plain, ((e1 = e2) | (~ spl30_22 | ~ spl30_23)), inference(backward_demodulation, [], [f465, f461])).
fof(f2053, plain, (~ spl30_2 | ~ spl30_50), inference(avatar_split_clause, [], [f2052, f578, f374])).
fof(f2052, plain, (~ (e1 = op(e3, e3)) | ~ spl30_50), inference(backward_demodulation, [], [f183, f580])).
fof(f1982, plain, (~ spl30_2 | ~ spl30_4), inference(avatar_contradiction_clause, [], [f1981])).
fof(f1981, plain, ($false | (~ spl30_2 | ~ spl30_4)), inference(subsumption_resolution, [], [f1980, f215])).
fof(f1980, plain, ((e1 = e3) | (~ spl30_2 | ~ spl30_4)), inference(forward_demodulation, [], [f384, f376])).
fof(f1979, plain, (~ spl30_6 | ~ spl30_11 | spl30_149), inference(avatar_contradiction_clause, [], [f1978])).
fof(f1978, plain, ($false | (~ spl30_6 | ~ spl30_11 | spl30_149)), inference(subsumption_resolution, [], [f1973, f414])).
fof(f1973, plain, (~ (e2 = op(e3, e1)) | (~ spl30_6 | spl30_149)), inference(backward_demodulation, [], [f1086, f393])).
fof(f1977, plain, (~ spl30_6 | ~ spl30_33 | ~ spl30_148), inference(avatar_contradiction_clause, [], [f1976])).
fof(f1976, plain, ($false | (~ spl30_6 | ~ spl30_33 | ~ spl30_148)), inference(subsumption_resolution, [], [f1975, f212])).
fof(f1975, plain, ((e0 = e2) | (~ spl30_6 | ~ spl30_33 | ~ spl30_148)), inference(forward_demodulation, [], [f1972, f508])).
fof(f1972, plain, ((e2 = op(e1, e3)) | (~ spl30_6 | ~ spl30_148)), inference(backward_demodulation, [], [f1081, f393])).
fof(f1955, plain, (~ spl30_25 | ~ spl30_28), inference(avatar_contradiction_clause, [], [f1954])).
fof(f1954, plain, ($false | (~ spl30_25 | ~ spl30_28)), inference(subsumption_resolution, [], [f1953, f213])).
fof(f1953, plain, ((e0 = e3) | (~ spl30_25 | ~ spl30_28)), inference(backward_demodulation, [], [f486, f474])).
fof(f1945, plain, (~ spl30_42 | ~ spl30_43), inference(avatar_contradiction_clause, [], [f1944])).
fof(f1944, plain, ($false | (~ spl30_42 | ~ spl30_43)), inference(subsumption_resolution, [], [f1943, f214])).
fof(f1943, plain, ((e1 = e2) | (~ spl30_42 | ~ spl30_43)), inference(backward_demodulation, [], [f550, f546])).
fof(f1930, plain, (~ spl30_19 | ~ spl30_51), inference(avatar_split_clause, [], [f1928, f582, f446])).
fof(f1928, plain, (~ (e2 = op(e2, e3)) | ~ spl30_51), inference(backward_demodulation, [], [f182, f584])).
fof(f1925, plain, (~ spl30_53 | ~ spl30_56), inference(avatar_contradiction_clause, [], [f1924])).
fof(f1924, plain, ($false | (~ spl30_53 | ~ spl30_56)), inference(subsumption_resolution, [], [f1922, f213])).
fof(f1922, plain, ((e0 = e3) | (~ spl30_53 | ~ spl30_56)), inference(backward_demodulation, [], [f605, f593])).
fof(f1916, plain, (~ spl30_44 | ~ spl30_60), inference(avatar_split_clause, [], [f1912, f620, f552])).
fof(f1912, plain, (~ (e3 = op(e1, e1)) | ~ spl30_60), inference(backward_demodulation, [], [f169, f622])).
fof(f1908, plain, (~ spl30_62 | ~ spl30_63), inference(avatar_contradiction_clause, [], [f1907])).
fof(f1907, plain, ($false | (~ spl30_62 | ~ spl30_63)), inference(subsumption_resolution, [], [f1906, f214])).
fof(f1906, plain, ((e1 = e2) | (~ spl30_62 | ~ spl30_63)), inference(forward_demodulation, [], [f635, f631])).
fof(f1859, plain, (~ spl30_13 | ~ spl30_16), inference(avatar_contradiction_clause, [], [f1858])).
fof(f1858, plain, ($false | (~ spl30_13 | ~ spl30_16)), inference(subsumption_resolution, [], [f1856, f213])).
fof(f1856, plain, ((e0 = e3) | (~ spl30_13 | ~ spl30_16)), inference(backward_demodulation, [], [f435, f423])).
fof(f1853, plain, (~ spl30_21 | ~ spl30_30 | spl30_220), inference(avatar_contradiction_clause, [], [f1852])).
fof(f1852, plain, ($false | (~ spl30_21 | ~ spl30_30 | spl30_220)), inference(subsumption_resolution, [], [f1850, f495])).
fof(f1850, plain, (~ (e1 = op(e2, e0)) | (~ spl30_21 | spl30_220)), inference(backward_demodulation, [], [f1434, f457])).
fof(f1845, plain, (~ spl30_30 | ~ spl30_31), inference(avatar_contradiction_clause, [], [f1844])).
fof(f1844, plain, ($false | (~ spl30_30 | ~ spl30_31)), inference(subsumption_resolution, [], [f1843, f214])).
fof(f1843, plain, ((e1 = e2) | (~ spl30_30 | ~ spl30_31)), inference(backward_demodulation, [], [f499, f495])).
fof(f1842, plain, (~ spl30_38 | ~ spl30_40), inference(avatar_contradiction_clause, [], [f1841])).
fof(f1841, plain, ($false | (~ spl30_38 | ~ spl30_40)), inference(subsumption_resolution, [], [f1840, f215])).
fof(f1840, plain, ((e1 = e3) | (~ spl30_38 | ~ spl30_40)), inference(backward_demodulation, [], [f537, f529])).
fof(f1838, plain, (~ spl30_35 | ~ spl30_41 | ~ spl30_48 | spl30_213), inference(avatar_contradiction_clause, [], [f1837])).
fof(f1837, plain, ($false | (~ spl30_35 | ~ spl30_41 | ~ spl30_48 | spl30_213)), inference(subsumption_resolution, [], [f1835, f516])).
fof(f1835, plain, (~ (e2 = op(e1, e3)) | (~ spl30_41 | ~ spl30_48 | spl30_213)), inference(backward_demodulation, [], [f1540, f571])).
fof(f1540, plain, (~ (e2 = op(e1, op(e1, e0))) | (~ spl30_41 | spl30_213)), inference(backward_demodulation, [], [f1399, f542])).
fof(f1832, plain, (~ spl30_49 | ~ spl30_52), inference(avatar_contradiction_clause, [], [f1831])).
fof(f1831, plain, ($false | (~ spl30_49 | ~ spl30_52)), inference(subsumption_resolution, [], [f1829, f213])).
fof(f1829, plain, ((e0 = e3) | (~ spl30_49 | ~ spl30_52)), inference(backward_demodulation, [], [f588, f576])).
fof(f1821, plain, (~ spl30_46 | ~ spl30_62), inference(avatar_split_clause, [], [f1815, f629, f561])).
fof(f1815, plain, (~ (e1 = op(e1, e0)) | ~ spl30_62), inference(backward_demodulation, [], [f163, f631])).
fof(f1781, plain, (~ spl30_7 | spl30_149), inference(avatar_contradiction_clause, [], [f1780])).
fof(f1780, plain, ($false | (~ spl30_7 | spl30_149)), inference(subsumption_resolution, [], [f1775, f397])).
fof(f1775, plain, (~ (e2 = op(e3, e2)) | (~ spl30_7 | spl30_149)), inference(backward_demodulation, [], [f1086, f397])).
fof(f1774, plain, (~ spl30_5 | ~ spl30_2 | ~ spl30_11 | spl30_227), inference(avatar_split_clause, [], [f1673, f1467, f412, f374, f387])).
fof(f1673, plain, (~ (e0 = op(e3, e2)) | (~ spl30_2 | ~ spl30_11 | spl30_227)), inference(backward_demodulation, [], [f1650, f414])).
fof(f1650, plain, (~ (e0 = op(e3, op(e3, e1))) | (~ spl30_2 | spl30_227)), inference(backward_demodulation, [], [f1469, f376])).
fof(f1732, plain, (~ spl30_26 | ~ spl30_58), inference(avatar_split_clause, [], [f1730, f612, f476])).
fof(f1730, plain, (~ (e1 = op(e2, e1)) | ~ spl30_58), inference(backward_demodulation, [], [f170, f614])).
fof(f1723, plain, (~ spl30_57 | ~ spl30_2 | spl30_203), inference(avatar_split_clause, [], [f1648, f1346, f374, f608])).
fof(f1648, plain, (~ (e0 = op(e0, e1)) | (~ spl30_2 | spl30_203)), inference(backward_demodulation, [], [f1348, f376])).
fof(f1718, plain, (~ spl30_50 | ~ spl30_52), inference(avatar_contradiction_clause, [], [f1717])).
fof(f1717, plain, ($false | (~ spl30_50 | ~ spl30_52)), inference(subsumption_resolution, [], [f1716, f215])).
fof(f1716, plain, ((e1 = e3) | (~ spl30_50 | ~ spl30_52)), inference(forward_demodulation, [], [f588, f580])).
fof(f1715, plain, (~ spl30_46 | ~ spl30_14), inference(avatar_split_clause, [], [f1714, f425, f561])).
fof(f1714, plain, (~ (e1 = op(e1, e0)) | ~ spl30_14), inference(forward_demodulation, [], [f167, f427])).
fof(f1713, plain, (~ spl30_45 | ~ spl30_41), inference(avatar_split_clause, [], [f1712, f540, f557])).
fof(f1712, plain, (~ (e0 = op(e1, e0)) | ~ spl30_41), inference(forward_demodulation, [], [f193, f542])).
fof(f1696, plain, (~ spl30_27 | ~ spl30_11), inference(avatar_split_clause, [], [f1695, f412, f480])).
fof(f1695, plain, (~ (e2 = op(e2, e1)) | ~ spl30_11), inference(forward_demodulation, [], [f174, f414])).
fof(f1685, plain, (~ spl30_17 | ~ spl30_19), inference(avatar_contradiction_clause, [], [f1684])).
fof(f1684, plain, ($false | (~ spl30_17 | ~ spl30_19)), inference(subsumption_resolution, [], [f1683, f212])).
fof(f1683, plain, ((e0 = e2) | (~ spl30_17 | ~ spl30_19)), inference(forward_demodulation, [], [f448, f440])).
fof(f1682, plain, (~ spl30_14 | ~ spl30_16), inference(avatar_contradiction_clause, [], [f1681])).
fof(f1681, plain, ($false | (~ spl30_14 | ~ spl30_16)), inference(subsumption_resolution, [], [f1680, f215])).
fof(f1680, plain, ((e1 = e3) | (~ spl30_14 | ~ spl30_16)), inference(forward_demodulation, [], [f435, f427])).
fof(f1672, plain, (~ spl30_12 | ~ spl30_2 | spl30_200), inference(avatar_split_clause, [], [f1645, f1331, f374, f416])).
fof(f1645, plain, (~ (e3 = op(e3, e1)) | (~ spl30_2 | spl30_200)), inference(backward_demodulation, [], [f1333, f376])).
fof(f1661, plain, (~ spl30_5 | ~ spl30_7), inference(avatar_contradiction_clause, [], [f1660])).
fof(f1660, plain, ($false | (~ spl30_5 | ~ spl30_7)), inference(subsumption_resolution, [], [f1659, f212])).
fof(f1659, plain, ((e0 = e2) | (~ spl30_5 | ~ spl30_7)), inference(forward_demodulation, [], [f397, f389])).
fof(f1644, plain, (~ spl30_5 | ~ spl30_6), inference(avatar_contradiction_clause, [], [f1643])).
fof(f1643, plain, ($false | (~ spl30_5 | ~ spl30_6)), inference(subsumption_resolution, [], [f1642, f211])).
fof(f1642, plain, ((e0 = e1) | (~ spl30_5 | ~ spl30_6)), inference(backward_demodulation, [], [f393, f389])).
fof(f1641, plain, (~ spl30_6 | ~ spl30_7), inference(avatar_contradiction_clause, [], [f1640])).
fof(f1640, plain, ($false | (~ spl30_6 | ~ spl30_7)), inference(subsumption_resolution, [], [f1639, f214])).
fof(f1639, plain, ((e1 = e2) | (~ spl30_6 | ~ spl30_7)), inference(backward_demodulation, [], [f397, f393])).
fof(f1638, plain, (~ spl30_7 | ~ spl30_8), inference(avatar_contradiction_clause, [], [f1637])).
fof(f1637, plain, ($false | (~ spl30_7 | ~ spl30_8)), inference(subsumption_resolution, [], [f1636, f216])).
fof(f1636, plain, ((e2 = e3) | (~ spl30_7 | ~ spl30_8)), inference(backward_demodulation, [], [f401, f397])).
fof(f1633, plain, (~ spl30_8 | spl30_147), inference(avatar_contradiction_clause, [], [f1632])).
fof(f1632, plain, ($false | (~ spl30_8 | spl30_147)), inference(subsumption_resolution, [], [f1628, f401])).
fof(f1628, plain, (~ (e3 = op(e3, e2)) | (~ spl30_8 | spl30_147)), inference(backward_demodulation, [], [f1076, f401])).
fof(f1626, plain, (~ spl30_10 | ~ spl30_11), inference(avatar_contradiction_clause, [], [f1625])).
fof(f1625, plain, ($false | (~ spl30_10 | ~ spl30_11)), inference(subsumption_resolution, [], [f1624, f214])).
fof(f1624, plain, ((e1 = e2) | (~ spl30_10 | ~ spl30_11)), inference(backward_demodulation, [], [f414, f410])).
fof(f1623, plain, (~ spl30_11 | ~ spl30_12), inference(avatar_contradiction_clause, [], [f1622])).
fof(f1622, plain, ($false | (~ spl30_11 | ~ spl30_12)), inference(subsumption_resolution, [], [f1621, f216])).
fof(f1621, plain, ((e2 = e3) | (~ spl30_11 | ~ spl30_12)), inference(backward_demodulation, [], [f418, f414])).
fof(f1616, plain, (~ spl30_14 | ~ spl30_15), inference(avatar_contradiction_clause, [], [f1615])).
fof(f1615, plain, ($false | (~ spl30_14 | ~ spl30_15)), inference(subsumption_resolution, [], [f1614, f214])).
fof(f1614, plain, ((e1 = e2) | (~ spl30_14 | ~ spl30_15)), inference(backward_demodulation, [], [f431, f427])).
fof(f1613, plain, (~ spl30_15 | ~ spl30_16), inference(avatar_contradiction_clause, [], [f1612])).
fof(f1612, plain, ($false | (~ spl30_15 | ~ spl30_16)), inference(subsumption_resolution, [], [f1611, f216])).
fof(f1611, plain, ((e2 = e3) | (~ spl30_15 | ~ spl30_16)), inference(backward_demodulation, [], [f435, f431])).
fof(f1604, plain, (~ spl30_17 | ~ spl30_18), inference(avatar_contradiction_clause, [], [f1603])).
fof(f1603, plain, ($false | (~ spl30_17 | ~ spl30_18)), inference(subsumption_resolution, [], [f1602, f211])).
fof(f1602, plain, ((e0 = e1) | (~ spl30_17 | ~ spl30_18)), inference(backward_demodulation, [], [f444, f440])).
fof(f1601, plain, (~ spl30_18 | ~ spl30_19), inference(avatar_contradiction_clause, [], [f1600])).
fof(f1600, plain, ($false | (~ spl30_18 | ~ spl30_19)), inference(subsumption_resolution, [], [f1599, f214])).
fof(f1599, plain, ((e1 = e2) | (~ spl30_18 | ~ spl30_19)), inference(backward_demodulation, [], [f448, f444])).
fof(f1598, plain, (~ spl30_19 | ~ spl30_20), inference(avatar_contradiction_clause, [], [f1597])).
fof(f1597, plain, ($false | (~ spl30_19 | ~ spl30_20)), inference(subsumption_resolution, [], [f1596, f216])).
fof(f1596, plain, ((e2 = e3) | (~ spl30_19 | ~ spl30_20)), inference(backward_demodulation, [], [f452, f448])).
fof(f1581, plain, (~ spl30_21 | ~ spl30_25), inference(avatar_split_clause, [], [f1578, f472, f455])).
fof(f1578, plain, (~ (e0 = op(e2, e2)) | ~ spl30_25), inference(backward_demodulation, [], [f202, f474])).
fof(f1568, plain, (~ spl30_34 | ~ spl30_35), inference(avatar_contradiction_clause, [], [f1567])).
fof(f1567, plain, ($false | (~ spl30_34 | ~ spl30_35)), inference(subsumption_resolution, [], [f1566, f214])).
fof(f1566, plain, ((e1 = e2) | (~ spl30_34 | ~ spl30_35)), inference(backward_demodulation, [], [f516, f512])).
fof(f1565, plain, (~ spl30_35 | ~ spl30_36), inference(avatar_contradiction_clause, [], [f1564])).
fof(f1564, plain, ($false | (~ spl30_35 | ~ spl30_36)), inference(subsumption_resolution, [], [f1563, f216])).
fof(f1563, plain, ((e2 = e3) | (~ spl30_35 | ~ spl30_36)), inference(backward_demodulation, [], [f520, f516])).
fof(f1531, plain, (~ spl30_13 | ~ spl30_45), inference(avatar_split_clause, [], [f1526, f557, f421])).
fof(f1526, plain, (~ (e0 = op(e3, e0)) | ~ spl30_45), inference(backward_demodulation, [], [f167, f559])).
fof(f1524, plain, (~ spl30_50 | ~ spl30_51), inference(avatar_contradiction_clause, [], [f1523])).
fof(f1523, plain, ($false | (~ spl30_50 | ~ spl30_51)), inference(subsumption_resolution, [], [f1522, f214])).
fof(f1522, plain, ((e1 = e2) | (~ spl30_50 | ~ spl30_51)), inference(backward_demodulation, [], [f584, f580])).
fof(f1521, plain, (~ spl30_51 | ~ spl30_52), inference(avatar_contradiction_clause, [], [f1520])).
fof(f1520, plain, ($false | (~ spl30_51 | ~ spl30_52)), inference(subsumption_resolution, [], [f1519, f216])).
fof(f1519, plain, ((e2 = e3) | (~ spl30_51 | ~ spl30_52)), inference(backward_demodulation, [], [f588, f584])).
fof(f1501, plain, (~ spl30_25 | ~ spl30_57), inference(avatar_split_clause, [], [f1496, f608, f472])).
fof(f1496, plain, (~ (e0 = op(e2, e1)) | ~ spl30_57), inference(backward_demodulation, [], [f170, f610])).
fof(f1476, plain, (~ spl30_228 | ~ spl30_2 | ~ spl30_227), inference(avatar_split_clause, [], [f368, f1467, f374, f1472])).
fof(f368, plain, (~ (e0 = op(e3, op(e3, op(e3, e3)))) | ~ (e1 = op(e3, e3)) | ~ (e2 = op(e3, op(e3, e3)))), inference(cnf_transformation, [], [f53])).
fof(f53, plain, (~ (e0 = op(e3, op(e3, op(e3, e3)))) | ~ (e1 = op(e3, e3)) | ~ (e2 = op(e3, op(e3, e3)))), inference(ennf_transformation, [], [f29])).
fof(f29, plain, ~ ((e0 = op(e3, op(e3, op(e3, e3)))) & (e1 = op(e3, e3)) & (e2 = op(e3, op(e3, e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG157+1.p', ax29)).
fof(f1475, plain, (~ spl30_228 | ~ spl30_1 | ~ spl30_225), inference(avatar_split_clause, [], [f367, f1457, f370, f1472])).
fof(f367, plain, (~ (e1 = op(e3, op(e3, op(e3, e3)))) | ~ (e0 = op(e3, e3)) | ~ (e2 = op(e3, op(e3, e3)))), inference(cnf_transformation, [], [f52])).
fof(f52, plain, (~ (e1 = op(e3, op(e3, op(e3, e3)))) | ~ (e0 = op(e3, e3)) | ~ (e2 = op(e3, op(e3, e3)))), inference(ennf_transformation, [], [f28])).
fof(f28, plain, ~ ((e1 = op(e3, op(e3, op(e3, e3)))) & (e0 = op(e3, e3)) & (e2 = op(e3, op(e3, e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG157+1.p', ax28)).
fof(f1470, plain, (~ spl30_226 | ~ spl30_3 | ~ spl30_227), inference(avatar_split_clause, [], [f366, f1467, f378, f1462])).
fof(f366, plain, (~ (e0 = op(e3, op(e3, op(e3, e3)))) | ~ (e2 = op(e3, e3)) | ~ (e1 = op(e3, op(e3, e3)))), inference(cnf_transformation, [], [f51])).
fof(f51, plain, (~ (e0 = op(e3, op(e3, op(e3, e3)))) | ~ (e2 = op(e3, e3)) | ~ (e1 = op(e3, op(e3, e3)))), inference(ennf_transformation, [], [f27])).
fof(f27, plain, ~ ((e0 = op(e3, op(e3, op(e3, e3)))) & (e2 = op(e3, e3)) & (e1 = op(e3, op(e3, e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG157+1.p', ax27)).
fof(f1465, plain, (~ spl30_226 | ~ spl30_1 | ~ spl30_224), inference(avatar_split_clause, [], [f365, f1452, f370, f1462])).
fof(f365, plain, (~ (e2 = op(e3, op(e3, op(e3, e3)))) | ~ (e0 = op(e3, e3)) | ~ (e1 = op(e3, op(e3, e3)))), inference(cnf_transformation, [], [f50])).
fof(f50, plain, (~ (e2 = op(e3, op(e3, op(e3, e3)))) | ~ (e0 = op(e3, e3)) | ~ (e1 = op(e3, op(e3, e3)))), inference(ennf_transformation, [], [f26])).
fof(f26, plain, ~ ((e2 = op(e3, op(e3, op(e3, e3)))) & (e0 = op(e3, e3)) & (e1 = op(e3, op(e3, e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG157+1.p', ax26)).
fof(f1460, plain, (~ spl30_223 | ~ spl30_3 | ~ spl30_225), inference(avatar_split_clause, [], [f364, f1457, f378, f1448])).
fof(f364, plain, (~ (e1 = op(e3, op(e3, op(e3, e3)))) | ~ (e2 = op(e3, e3)) | ~ (e0 = op(e3, op(e3, e3)))), inference(cnf_transformation, [], [f49])).
fof(f49, plain, (~ (e1 = op(e3, op(e3, op(e3, e3)))) | ~ (e2 = op(e3, e3)) | ~ (e0 = op(e3, op(e3, e3)))), inference(ennf_transformation, [], [f25])).
fof(f25, plain, ~ ((e1 = op(e3, op(e3, op(e3, e3)))) & (e2 = op(e3, e3)) & (e0 = op(e3, op(e3, e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG157+1.p', ax25)).
fof(f1455, plain, (~ spl30_223 | ~ spl30_2 | ~ spl30_224), inference(avatar_split_clause, [], [f363, f1452, f374, f1448])).
fof(f363, plain, (~ (e2 = op(e3, op(e3, op(e3, e3)))) | ~ (e1 = op(e3, e3)) | ~ (e0 = op(e3, op(e3, e3)))), inference(cnf_transformation, [], [f48])).
fof(f48, plain, (~ (e2 = op(e3, op(e3, op(e3, e3)))) | ~ (e1 = op(e3, e3)) | ~ (e0 = op(e3, op(e3, e3)))), inference(ennf_transformation, [], [f24])).
fof(f24, plain, ~ ((e2 = op(e3, op(e3, op(e3, e3)))) & (e1 = op(e3, e3)) & (e0 = op(e3, op(e3, e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG157+1.p', ax24)).
fof(f1446, plain, (~ spl30_222 | ~ spl30_22 | ~ spl30_221), inference(avatar_split_clause, [], [f362, f1437, f459, f1442])).
fof(f362, plain, (~ (e0 = op(e2, op(e2, op(e2, e2)))) | ~ (e1 = op(e2, e2)) | ~ (e3 = op(e2, op(e2, e2)))), inference(cnf_transformation, [], [f47])).
fof(f47, plain, (~ (e0 = op(e2, op(e2, op(e2, e2)))) | ~ (e1 = op(e2, e2)) | ~ (e3 = op(e2, op(e2, e2)))), inference(ennf_transformation, [], [f23])).
fof(f23, plain, ~ ((e0 = op(e2, op(e2, op(e2, e2)))) & (e1 = op(e2, e2)) & (e3 = op(e2, op(e2, e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG157+1.p', ax23)).
fof(f1445, plain, (~ spl30_222 | ~ spl30_21 | ~ spl30_219), inference(avatar_split_clause, [], [f361, f1427, f455, f1442])).
fof(f361, plain, (~ (e1 = op(e2, op(e2, op(e2, e2)))) | ~ (e0 = op(e2, e2)) | ~ (e3 = op(e2, op(e2, e2)))), inference(cnf_transformation, [], [f46])).
fof(f46, plain, (~ (e1 = op(e2, op(e2, op(e2, e2)))) | ~ (e0 = op(e2, e2)) | ~ (e3 = op(e2, op(e2, e2)))), inference(ennf_transformation, [], [f22])).
fof(f22, plain, ~ ((e1 = op(e2, op(e2, op(e2, e2)))) & (e0 = op(e2, e2)) & (e3 = op(e2, op(e2, e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG157+1.p', ax22)).
fof(f1440, plain, (~ spl30_220 | ~ spl30_24 | ~ spl30_221), inference(avatar_split_clause, [], [f360, f1437, f467, f1432])).
fof(f360, plain, (~ (e0 = op(e2, op(e2, op(e2, e2)))) | ~ (e3 = op(e2, e2)) | ~ (e1 = op(e2, op(e2, e2)))), inference(cnf_transformation, [], [f45])).
fof(f45, plain, (~ (e0 = op(e2, op(e2, op(e2, e2)))) | ~ (e3 = op(e2, e2)) | ~ (e1 = op(e2, op(e2, e2)))), inference(ennf_transformation, [], [f21])).
fof(f21, plain, ~ ((e0 = op(e2, op(e2, op(e2, e2)))) & (e3 = op(e2, e2)) & (e1 = op(e2, op(e2, e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG157+1.p', ax21)).
fof(f1435, plain, (~ spl30_220 | ~ spl30_21 | ~ spl30_218), inference(avatar_split_clause, [], [f359, f1422, f455, f1432])).
fof(f359, plain, (~ (e3 = op(e2, op(e2, op(e2, e2)))) | ~ (e0 = op(e2, e2)) | ~ (e1 = op(e2, op(e2, e2)))), inference(cnf_transformation, [], [f44])).
fof(f44, plain, (~ (e3 = op(e2, op(e2, op(e2, e2)))) | ~ (e0 = op(e2, e2)) | ~ (e1 = op(e2, op(e2, e2)))), inference(ennf_transformation, [], [f20])).
fof(f20, plain, ~ ((e3 = op(e2, op(e2, op(e2, e2)))) & (e0 = op(e2, e2)) & (e1 = op(e2, op(e2, e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG157+1.p', ax20)).
fof(f1430, plain, (~ spl30_217 | ~ spl30_24 | ~ spl30_219), inference(avatar_split_clause, [], [f358, f1427, f467, f1418])).
fof(f358, plain, (~ (e1 = op(e2, op(e2, op(e2, e2)))) | ~ (e3 = op(e2, e2)) | ~ (e0 = op(e2, op(e2, e2)))), inference(cnf_transformation, [], [f43])).
fof(f43, plain, (~ (e1 = op(e2, op(e2, op(e2, e2)))) | ~ (e3 = op(e2, e2)) | ~ (e0 = op(e2, op(e2, e2)))), inference(ennf_transformation, [], [f19])).
fof(f19, plain, ~ ((e1 = op(e2, op(e2, op(e2, e2)))) & (e3 = op(e2, e2)) & (e0 = op(e2, op(e2, e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG157+1.p', ax19)).
fof(f1425, plain, (~ spl30_217 | ~ spl30_22 | ~ spl30_218), inference(avatar_split_clause, [], [f357, f1422, f459, f1418])).
fof(f357, plain, (~ (e3 = op(e2, op(e2, op(e2, e2)))) | ~ (e1 = op(e2, e2)) | ~ (e0 = op(e2, op(e2, e2)))), inference(cnf_transformation, [], [f42])).
fof(f42, plain, (~ (e3 = op(e2, op(e2, op(e2, e2)))) | ~ (e1 = op(e2, e2)) | ~ (e0 = op(e2, op(e2, e2)))), inference(ennf_transformation, [], [f18])).
fof(f18, plain, ~ ((e3 = op(e2, op(e2, op(e2, e2)))) & (e1 = op(e2, e2)) & (e0 = op(e2, op(e2, e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG157+1.p', ax18)).
fof(f1416, plain, (~ spl30_216 | ~ spl30_43 | ~ spl30_215), inference(avatar_split_clause, [], [f356, f1407, f548, f1412])).
fof(f356, plain, (~ (e0 = op(e1, op(e1, op(e1, e1)))) | ~ (e2 = op(e1, e1)) | ~ (e3 = op(e1, op(e1, e1)))), inference(cnf_transformation, [], [f41])).
fof(f41, plain, (~ (e0 = op(e1, op(e1, op(e1, e1)))) | ~ (e2 = op(e1, e1)) | ~ (e3 = op(e1, op(e1, e1)))), inference(ennf_transformation, [], [f17])).
fof(f17, plain, ~ ((e0 = op(e1, op(e1, op(e1, e1)))) & (e2 = op(e1, e1)) & (e3 = op(e1, op(e1, e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG157+1.p', ax17)).
fof(f1415, plain, (~ spl30_216 | ~ spl30_41 | ~ spl30_213), inference(avatar_split_clause, [], [f355, f1397, f540, f1412])).
fof(f355, plain, (~ (e2 = op(e1, op(e1, op(e1, e1)))) | ~ (e0 = op(e1, e1)) | ~ (e3 = op(e1, op(e1, e1)))), inference(cnf_transformation, [], [f40])).
fof(f40, plain, (~ (e2 = op(e1, op(e1, op(e1, e1)))) | ~ (e0 = op(e1, e1)) | ~ (e3 = op(e1, op(e1, e1)))), inference(ennf_transformation, [], [f16])).
fof(f16, plain, ~ ((e2 = op(e1, op(e1, op(e1, e1)))) & (e0 = op(e1, e1)) & (e3 = op(e1, op(e1, e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG157+1.p', ax16)).
fof(f1410, plain, (~ spl30_214 | ~ spl30_44 | ~ spl30_215), inference(avatar_split_clause, [], [f354, f1407, f552, f1402])).
fof(f354, plain, (~ (e0 = op(e1, op(e1, op(e1, e1)))) | ~ (e3 = op(e1, e1)) | ~ (e2 = op(e1, op(e1, e1)))), inference(cnf_transformation, [], [f39])).
fof(f39, plain, (~ (e0 = op(e1, op(e1, op(e1, e1)))) | ~ (e3 = op(e1, e1)) | ~ (e2 = op(e1, op(e1, e1)))), inference(ennf_transformation, [], [f15])).
fof(f15, plain, ~ ((e0 = op(e1, op(e1, op(e1, e1)))) & (e3 = op(e1, e1)) & (e2 = op(e1, op(e1, e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG157+1.p', ax15)).
fof(f1405, plain, (~ spl30_214 | ~ spl30_41 | ~ spl30_212), inference(avatar_split_clause, [], [f353, f1392, f540, f1402])).
fof(f353, plain, (~ (e3 = op(e1, op(e1, op(e1, e1)))) | ~ (e0 = op(e1, e1)) | ~ (e2 = op(e1, op(e1, e1)))), inference(cnf_transformation, [], [f38])).
fof(f38, plain, (~ (e3 = op(e1, op(e1, op(e1, e1)))) | ~ (e0 = op(e1, e1)) | ~ (e2 = op(e1, op(e1, e1)))), inference(ennf_transformation, [], [f14])).
fof(f14, plain, ~ ((e3 = op(e1, op(e1, op(e1, e1)))) & (e0 = op(e1, e1)) & (e2 = op(e1, op(e1, e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG157+1.p', ax14)).
fof(f1400, plain, (~ spl30_211 | ~ spl30_44 | ~ spl30_213), inference(avatar_split_clause, [], [f352, f1397, f552, f1388])).
fof(f352, plain, (~ (e2 = op(e1, op(e1, op(e1, e1)))) | ~ (e3 = op(e1, e1)) | ~ (e0 = op(e1, op(e1, e1)))), inference(cnf_transformation, [], [f37])).
fof(f37, plain, (~ (e2 = op(e1, op(e1, op(e1, e1)))) | ~ (e3 = op(e1, e1)) | ~ (e0 = op(e1, op(e1, e1)))), inference(ennf_transformation, [], [f13])).
fof(f13, plain, ~ ((e2 = op(e1, op(e1, op(e1, e1)))) & (e3 = op(e1, e1)) & (e0 = op(e1, op(e1, e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG157+1.p', ax13)).
fof(f1395, plain, (~ spl30_211 | ~ spl30_43 | ~ spl30_212), inference(avatar_split_clause, [], [f351, f1392, f548, f1388])).
fof(f351, plain, (~ (e3 = op(e1, op(e1, op(e1, e1)))) | ~ (e2 = op(e1, e1)) | ~ (e0 = op(e1, op(e1, e1)))), inference(cnf_transformation, [], [f36])).
fof(f36, plain, (~ (e3 = op(e1, op(e1, op(e1, e1)))) | ~ (e2 = op(e1, e1)) | ~ (e0 = op(e1, op(e1, e1)))), inference(ennf_transformation, [], [f12])).
fof(f12, plain, ~ ((e3 = op(e1, op(e1, op(e1, e1)))) & (e2 = op(e1, e1)) & (e0 = op(e1, op(e1, e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG157+1.p', ax12)).
fof(f1386, plain, (~ spl30_210 | ~ spl30_63 | ~ spl30_209), inference(avatar_split_clause, [], [f350, f1377, f633, f1382])).
fof(f350, plain, (~ (e1 = op(e0, op(e0, op(e0, e0)))) | ~ (op(e0, e0) = e2) | ~ (e3 = op(e0, op(e0, e0)))), inference(cnf_transformation, [], [f35])).
fof(f35, plain, (~ (e1 = op(e0, op(e0, op(e0, e0)))) | ~ (op(e0, e0) = e2) | ~ (e3 = op(e0, op(e0, e0)))), inference(ennf_transformation, [], [f11])).
fof(f11, plain, ~ ((e1 = op(e0, op(e0, op(e0, e0)))) & (op(e0, e0) = e2) & (e3 = op(e0, op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG157+1.p', ax11)).
fof(f1385, plain, (~ spl30_210 | ~ spl30_62 | ~ spl30_207), inference(avatar_split_clause, [], [f349, f1367, f629, f1382])).
fof(f349, plain, (~ (e2 = op(e0, op(e0, op(e0, e0)))) | ~ (op(e0, e0) = e1) | ~ (e3 = op(e0, op(e0, e0)))), inference(cnf_transformation, [], [f34])).
fof(f34, plain, (~ (e2 = op(e0, op(e0, op(e0, e0)))) | ~ (op(e0, e0) = e1) | ~ (e3 = op(e0, op(e0, e0)))), inference(ennf_transformation, [], [f10])).
fof(f10, plain, ~ ((e2 = op(e0, op(e0, op(e0, e0)))) & (op(e0, e0) = e1) & (e3 = op(e0, op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG157+1.p', ax10)).
fof(f1380, plain, (~ spl30_208 | ~ spl30_64 | ~ spl30_209), inference(avatar_split_clause, [], [f348, f1377, f637, f1372])).
fof(f348, plain, (~ (e1 = op(e0, op(e0, op(e0, e0)))) | ~ (op(e0, e0) = e3) | ~ (e2 = op(e0, op(e0, e0)))), inference(cnf_transformation, [], [f33])).
fof(f33, plain, (~ (e1 = op(e0, op(e0, op(e0, e0)))) | ~ (op(e0, e0) = e3) | ~ (e2 = op(e0, op(e0, e0)))), inference(ennf_transformation, [], [f9])).
fof(f9, plain, ~ ((e1 = op(e0, op(e0, op(e0, e0)))) & (op(e0, e0) = e3) & (e2 = op(e0, op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG157+1.p', ax9)).
fof(f1375, plain, (~ spl30_208 | ~ spl30_62 | ~ spl30_206), inference(avatar_split_clause, [], [f347, f1362, f629, f1372])).
fof(f347, plain, (~ (e3 = op(e0, op(e0, op(e0, e0)))) | ~ (op(e0, e0) = e1) | ~ (e2 = op(e0, op(e0, e0)))), inference(cnf_transformation, [], [f32])).
fof(f32, plain, (~ (e3 = op(e0, op(e0, op(e0, e0)))) | ~ (op(e0, e0) = e1) | ~ (e2 = op(e0, op(e0, e0)))), inference(ennf_transformation, [], [f8])).
fof(f8, plain, ~ ((e3 = op(e0, op(e0, op(e0, e0)))) & (op(e0, e0) = e1) & (e2 = op(e0, op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG157+1.p', ax8)).
fof(f1370, plain, (~ spl30_205 | ~ spl30_64 | ~ spl30_207), inference(avatar_split_clause, [], [f346, f1367, f637, f1358])).
fof(f346, plain, (~ (e2 = op(e0, op(e0, op(e0, e0)))) | ~ (op(e0, e0) = e3) | ~ (e1 = op(e0, op(e0, e0)))), inference(cnf_transformation, [], [f31])).
fof(f31, plain, (~ (e2 = op(e0, op(e0, op(e0, e0)))) | ~ (op(e0, e0) = e3) | ~ (e1 = op(e0, op(e0, e0)))), inference(ennf_transformation, [], [f7])).
fof(f7, plain, ~ ((e2 = op(e0, op(e0, op(e0, e0)))) & (op(e0, e0) = e3) & (e1 = op(e0, op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG157+1.p', ax7)).
fof(f1365, plain, (~ spl30_205 | ~ spl30_63 | ~ spl30_206), inference(avatar_split_clause, [], [f345, f1362, f633, f1358])).
fof(f345, plain, (~ (e3 = op(e0, op(e0, op(e0, e0)))) | ~ (op(e0, e0) = e2) | ~ (e1 = op(e0, op(e0, e0)))), inference(cnf_transformation, [], [f30])).
fof(f30, plain, (~ (e3 = op(e0, op(e0, op(e0, e0)))) | ~ (op(e0, e0) = e2) | ~ (e1 = op(e0, op(e0, e0)))), inference(ennf_transformation, [], [f6])).
fof(f6, plain, ~ ((e3 = op(e0, op(e0, op(e0, e0)))) & (op(e0, e0) = e2) & (e1 = op(e0, op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG157+1.p', ax6)).
fof(f1356, plain, (spl30_198 | spl30_194 | spl30_190 | spl30_186 | spl30_182 | spl30_180 | spl30_176 | spl30_172 | spl30_168 | spl30_164 | spl30_162 | spl30_158 | spl30_154 | spl30_150 | spl30_146 | ~ spl30_200), inference(avatar_split_clause, [], [f337, f1331, f1070, f1089, f1108, f1127, f1146, f1157, f1176, f1195, f1214, f1233, f1244, f1263, f1282, f1301, f1320])).
fof(f1320, plain, (spl30_198 <=> sP0), introduced(avatar_definition, [new_symbols(naming, [spl30_198])])).
fof(f1301, plain, (spl30_194 <=> sP1), introduced(avatar_definition, [new_symbols(naming, [spl30_194])])).
fof(f1282, plain, (spl30_190 <=> sP2), introduced(avatar_definition, [new_symbols(naming, [spl30_190])])).
fof(f1263, plain, (spl30_186 <=> sP3), introduced(avatar_definition, [new_symbols(naming, [spl30_186])])).
fof(f1244, plain, (spl30_182 <=> sP4), introduced(avatar_definition, [new_symbols(naming, [spl30_182])])).
fof(f1233, plain, (spl30_180 <=> sP5), introduced(avatar_definition, [new_symbols(naming, [spl30_180])])).
fof(f1214, plain, (spl30_176 <=> sP6), introduced(avatar_definition, [new_symbols(naming, [spl30_176])])).
fof(f1195, plain, (spl30_172 <=> sP7), introduced(avatar_definition, [new_symbols(naming, [spl30_172])])).
fof(f1176, plain, (spl30_168 <=> sP8), introduced(avatar_definition, [new_symbols(naming, [spl30_168])])).
fof(f1157, plain, (spl30_164 <=> sP9), introduced(avatar_definition, [new_symbols(naming, [spl30_164])])).
fof(f1146, plain, (spl30_162 <=> sP10), introduced(avatar_definition, [new_symbols(naming, [spl30_162])])).
fof(f1127, plain, (spl30_158 <=> sP11), introduced(avatar_definition, [new_symbols(naming, [spl30_158])])).
fof(f1108, plain, (spl30_154 <=> sP12), introduced(avatar_definition, [new_symbols(naming, [spl30_154])])).
fof(f1089, plain, (spl30_150 <=> sP13), introduced(avatar_definition, [new_symbols(naming, [spl30_150])])).
fof(f1070, plain, (spl30_146 <=> sP14), introduced(avatar_definition, [new_symbols(naming, [spl30_146])])).
fof(f337, plain, (~ (e3 = op(e3, op(e3, e3))) | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f84])).
fof(f84, plain, (((~ (e3 = op(e3, op(e3, e3))) & ~ (e2 = op(e2, op(e3, e3))) & ~ (e1 = op(e1, op(e3, e3))) & ~ (e0 = op(e0, op(e3, e3))) & (op(e3, e3) = op(e3, e3))) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15) & ((~ (e3 = op(op(e3, e3), e3)) & (e3 = op(op(e3, e3), e3)) & ~ (e3 = op(e3, op(e3, e3)))) | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0)), inference(definition_folding, [], [f5, e83, e82, e81, e80, e79, e78, e77, e76, e75, e74, e73, e72, e71, e70, e69, e68, e67, e66, e65, e64, e63, e62, e61, e60, e59, e58, e57, e56, e55, e54])).
fof(f54, plain, ((~ (e0 = op(op(e0, e0), e0)) & (e0 = op(op(e0, e0), e0)) & ~ (e0 = op(e0, op(e0, e0)))) | ~ sP0), inference(usedef, [], [e54])).
fof(e54, plain, (sP0 <=> (~ (e0 = op(op(e0, e0), e0)) & (e0 = op(op(e0, e0), e0)) & ~ (e0 = op(e0, op(e0, e0))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f55, plain, ((~ (e0 = op(op(e0, e1), e1)) & (e1 = op(op(e0, e1), e0)) & ~ (e1 = op(e0, op(e0, e1)))) | ~ sP1), inference(usedef, [], [e55])).
fof(e55, plain, (sP1 <=> (~ (e0 = op(op(e0, e1), e1)) & (e1 = op(op(e0, e1), e0)) & ~ (e1 = op(e0, op(e0, e1))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f56, plain, ((~ (e0 = op(op(e0, e2), e2)) & (e2 = op(op(e0, e2), e0)) & ~ (e2 = op(e0, op(e0, e2)))) | ~ sP2), inference(usedef, [], [e56])).
fof(e56, plain, (sP2 <=> (~ (e0 = op(op(e0, e2), e2)) & (e2 = op(op(e0, e2), e0)) & ~ (e2 = op(e0, op(e0, e2))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f57, plain, ((~ (e0 = op(op(e0, e3), e3)) & (e3 = op(op(e0, e3), e0)) & ~ (e3 = op(e0, op(e0, e3)))) | ~ sP3), inference(usedef, [], [e57])).
fof(e57, plain, (sP3 <=> (~ (e0 = op(op(e0, e3), e3)) & (e3 = op(op(e0, e3), e0)) & ~ (e3 = op(e0, op(e0, e3))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f58, plain, ((~ (e1 = op(op(e1, e0), e0)) & (e0 = op(op(e1, e0), e1)) & ~ (e0 = op(e1, op(e1, e0)))) | ~ sP4), inference(usedef, [], [e58])).
fof(e58, plain, (sP4 <=> (~ (e1 = op(op(e1, e0), e0)) & (e0 = op(op(e1, e0), e1)) & ~ (e0 = op(e1, op(e1, e0))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f59, plain, ((~ (e1 = op(op(e1, e1), e1)) & (e1 = op(op(e1, e1), e1)) & ~ (e1 = op(e1, op(e1, e1)))) | ~ sP5), inference(usedef, [], [e59])).
fof(e59, plain, (sP5 <=> (~ (e1 = op(op(e1, e1), e1)) & (e1 = op(op(e1, e1), e1)) & ~ (e1 = op(e1, op(e1, e1))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f60, plain, ((~ (e1 = op(op(e1, e2), e2)) & (e2 = op(op(e1, e2), e1)) & ~ (e2 = op(e1, op(e1, e2)))) | ~ sP6), inference(usedef, [], [e60])).
fof(e60, plain, (sP6 <=> (~ (e1 = op(op(e1, e2), e2)) & (e2 = op(op(e1, e2), e1)) & ~ (e2 = op(e1, op(e1, e2))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f61, plain, ((~ (e1 = op(op(e1, e3), e3)) & (e3 = op(op(e1, e3), e1)) & ~ (e3 = op(e1, op(e1, e3)))) | ~ sP7), inference(usedef, [], [e61])).
fof(e61, plain, (sP7 <=> (~ (e1 = op(op(e1, e3), e3)) & (e3 = op(op(e1, e3), e1)) & ~ (e3 = op(e1, op(e1, e3))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f62, plain, ((~ (e2 = op(op(e2, e0), e0)) & (e0 = op(op(e2, e0), e2)) & ~ (e0 = op(e2, op(e2, e0)))) | ~ sP8), inference(usedef, [], [e62])).
fof(e62, plain, (sP8 <=> (~ (e2 = op(op(e2, e0), e0)) & (e0 = op(op(e2, e0), e2)) & ~ (e0 = op(e2, op(e2, e0))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f63, plain, ((~ (e2 = op(op(e2, e1), e1)) & (e1 = op(op(e2, e1), e2)) & ~ (e1 = op(e2, op(e2, e1)))) | ~ sP9), inference(usedef, [], [e63])).
fof(e63, plain, (sP9 <=> (~ (e2 = op(op(e2, e1), e1)) & (e1 = op(op(e2, e1), e2)) & ~ (e1 = op(e2, op(e2, e1))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f64, plain, ((~ (e2 = op(op(e2, e2), e2)) & (e2 = op(op(e2, e2), e2)) & ~ (e2 = op(e2, op(e2, e2)))) | ~ sP10), inference(usedef, [], [e64])).
fof(e64, plain, (sP10 <=> (~ (e2 = op(op(e2, e2), e2)) & (e2 = op(op(e2, e2), e2)) & ~ (e2 = op(e2, op(e2, e2))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f65, plain, ((~ (e2 = op(op(e2, e3), e3)) & (e3 = op(op(e2, e3), e2)) & ~ (e3 = op(e2, op(e2, e3)))) | ~ sP11), inference(usedef, [], [e65])).
fof(e65, plain, (sP11 <=> (~ (e2 = op(op(e2, e3), e3)) & (e3 = op(op(e2, e3), e2)) & ~ (e3 = op(e2, op(e2, e3))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f66, plain, ((~ (e3 = op(op(e3, e0), e0)) & (e0 = op(op(e3, e0), e3)) & ~ (e0 = op(e3, op(e3, e0)))) | ~ sP12), inference(usedef, [], [e66])).
fof(e66, plain, (sP12 <=> (~ (e3 = op(op(e3, e0), e0)) & (e0 = op(op(e3, e0), e3)) & ~ (e0 = op(e3, op(e3, e0))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f67, plain, ((~ (e3 = op(op(e3, e1), e1)) & (e1 = op(op(e3, e1), e3)) & ~ (e1 = op(e3, op(e3, e1)))) | ~ sP13), inference(usedef, [], [e67])).
fof(e67, plain, (sP13 <=> (~ (e3 = op(op(e3, e1), e1)) & (e1 = op(op(e3, e1), e3)) & ~ (e1 = op(e3, op(e3, e1))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f68, plain, ((~ (e3 = op(op(e3, e2), e2)) & (e2 = op(op(e3, e2), e3)) & ~ (e2 = op(e3, op(e3, e2)))) | ~ sP14), inference(usedef, [], [e68])).
fof(e68, plain, (sP14 <=> (~ (e3 = op(op(e3, e2), e2)) & (e2 = op(op(e3, e2), e3)) & ~ (e2 = op(e3, op(e3, e2))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f69, plain, ((~ (e3 = op(e3, op(e0, e0))) & ~ (e2 = op(e2, op(e0, e0))) & ~ (e1 = op(e1, op(e0, e0))) & ~ (e0 = op(e0, op(e0, e0))) & (op(e0, e0) = op(e0, e0))) | ~ sP15), inference(usedef, [], [e69])).
fof(e69, plain, (sP15 <=> (~ (e3 = op(e3, op(e0, e0))) & ~ (e2 = op(e2, op(e0, e0))) & ~ (e1 = op(e1, op(e0, e0))) & ~ (e0 = op(e0, op(e0, e0))) & (op(e0, e0) = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f70, plain, ((~ (e3 = op(e3, op(e0, e1))) & ~ (e2 = op(e2, op(e0, e1))) & ~ (e1 = op(e1, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e1))) & (op(e0, e1) = op(e1, e0))) | ~ sP16), inference(usedef, [], [e70])).
fof(e70, plain, (sP16 <=> (~ (e3 = op(e3, op(e0, e1))) & ~ (e2 = op(e2, op(e0, e1))) & ~ (e1 = op(e1, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e1))) & (op(e0, e1) = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f71, plain, ((~ (e3 = op(e3, op(e0, e2))) & ~ (e2 = op(e2, op(e0, e2))) & ~ (e1 = op(e1, op(e0, e2))) & ~ (e0 = op(e0, op(e0, e2))) & (op(e0, e2) = op(e2, e0))) | ~ sP17), inference(usedef, [], [e71])).
fof(e71, plain, (sP17 <=> (~ (e3 = op(e3, op(e0, e2))) & ~ (e2 = op(e2, op(e0, e2))) & ~ (e1 = op(e1, op(e0, e2))) & ~ (e0 = op(e0, op(e0, e2))) & (op(e0, e2) = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f72, plain, ((~ (e3 = op(e3, op(e0, e3))) & ~ (e2 = op(e2, op(e0, e3))) & ~ (e1 = op(e1, op(e0, e3))) & ~ (e0 = op(e0, op(e0, e3))) & (op(e0, e3) = op(e3, e0))) | ~ sP18), inference(usedef, [], [e72])).
fof(e72, plain, (sP18 <=> (~ (e3 = op(e3, op(e0, e3))) & ~ (e2 = op(e2, op(e0, e3))) & ~ (e1 = op(e1, op(e0, e3))) & ~ (e0 = op(e0, op(e0, e3))) & (op(e0, e3) = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f73, plain, ((~ (e3 = op(e3, op(e1, e0))) & ~ (e2 = op(e2, op(e1, e0))) & ~ (e1 = op(e1, op(e1, e0))) & ~ (e0 = op(e0, op(e1, e0))) & (op(e0, e1) = op(e1, e0))) | ~ sP19), inference(usedef, [], [e73])).
fof(e73, plain, (sP19 <=> (~ (e3 = op(e3, op(e1, e0))) & ~ (e2 = op(e2, op(e1, e0))) & ~ (e1 = op(e1, op(e1, e0))) & ~ (e0 = op(e0, op(e1, e0))) & (op(e0, e1) = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP19])])).
fof(f74, plain, ((~ (e3 = op(e3, op(e1, e1))) & ~ (e2 = op(e2, op(e1, e1))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e0, op(e1, e1))) & (op(e1, e1) = op(e1, e1))) | ~ sP20), inference(usedef, [], [e74])).
fof(e74, plain, (sP20 <=> (~ (e3 = op(e3, op(e1, e1))) & ~ (e2 = op(e2, op(e1, e1))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e0, op(e1, e1))) & (op(e1, e1) = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP20])])).
fof(f75, plain, ((~ (e3 = op(e3, op(e1, e2))) & ~ (e2 = op(e2, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e2))) & ~ (e0 = op(e0, op(e1, e2))) & (op(e1, e2) = op(e2, e1))) | ~ sP21), inference(usedef, [], [e75])).
fof(e75, plain, (sP21 <=> (~ (e3 = op(e3, op(e1, e2))) & ~ (e2 = op(e2, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e2))) & ~ (e0 = op(e0, op(e1, e2))) & (op(e1, e2) = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP21])])).
fof(f76, plain, ((~ (e3 = op(e3, op(e1, e3))) & ~ (e2 = op(e2, op(e1, e3))) & ~ (e1 = op(e1, op(e1, e3))) & ~ (e0 = op(e0, op(e1, e3))) & (op(e1, e3) = op(e3, e1))) | ~ sP22), inference(usedef, [], [e76])).
fof(e76, plain, (sP22 <=> (~ (e3 = op(e3, op(e1, e3))) & ~ (e2 = op(e2, op(e1, e3))) & ~ (e1 = op(e1, op(e1, e3))) & ~ (e0 = op(e0, op(e1, e3))) & (op(e1, e3) = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP22])])).
fof(f77, plain, ((~ (e3 = op(e3, op(e2, e0))) & ~ (e2 = op(e2, op(e2, e0))) & ~ (e1 = op(e1, op(e2, e0))) & ~ (e0 = op(e0, op(e2, e0))) & (op(e0, e2) = op(e2, e0))) | ~ sP23), inference(usedef, [], [e77])).
fof(e77, plain, (sP23 <=> (~ (e3 = op(e3, op(e2, e0))) & ~ (e2 = op(e2, op(e2, e0))) & ~ (e1 = op(e1, op(e2, e0))) & ~ (e0 = op(e0, op(e2, e0))) & (op(e0, e2) = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP23])])).
fof(f78, plain, ((~ (e3 = op(e3, op(e2, e1))) & ~ (e2 = op(e2, op(e2, e1))) & ~ (e1 = op(e1, op(e2, e1))) & ~ (e0 = op(e0, op(e2, e1))) & (op(e1, e2) = op(e2, e1))) | ~ sP24), inference(usedef, [], [e78])).
fof(e78, plain, (sP24 <=> (~ (e3 = op(e3, op(e2, e1))) & ~ (e2 = op(e2, op(e2, e1))) & ~ (e1 = op(e1, op(e2, e1))) & ~ (e0 = op(e0, op(e2, e1))) & (op(e1, e2) = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP24])])).
fof(f79, plain, ((~ (e3 = op(e3, op(e2, e2))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e1, op(e2, e2))) & ~ (e0 = op(e0, op(e2, e2))) & (op(e2, e2) = op(e2, e2))) | ~ sP25), inference(usedef, [], [e79])).
fof(e79, plain, (sP25 <=> (~ (e3 = op(e3, op(e2, e2))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e1, op(e2, e2))) & ~ (e0 = op(e0, op(e2, e2))) & (op(e2, e2) = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP25])])).
fof(f80, plain, ((~ (e3 = op(e3, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e3))) & ~ (e1 = op(e1, op(e2, e3))) & ~ (e0 = op(e0, op(e2, e3))) & (op(e2, e3) = op(e3, e2))) | ~ sP26), inference(usedef, [], [e80])).
fof(e80, plain, (sP26 <=> (~ (e3 = op(e3, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e3))) & ~ (e1 = op(e1, op(e2, e3))) & ~ (e0 = op(e0, op(e2, e3))) & (op(e2, e3) = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP26])])).
fof(f81, plain, ((~ (e3 = op(e3, op(e3, e0))) & ~ (e2 = op(e2, op(e3, e0))) & ~ (e1 = op(e1, op(e3, e0))) & ~ (e0 = op(e0, op(e3, e0))) & (op(e0, e3) = op(e3, e0))) | ~ sP27), inference(usedef, [], [e81])).
fof(e81, plain, (sP27 <=> (~ (e3 = op(e3, op(e3, e0))) & ~ (e2 = op(e2, op(e3, e0))) & ~ (e1 = op(e1, op(e3, e0))) & ~ (e0 = op(e0, op(e3, e0))) & (op(e0, e3) = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP27])])).
fof(f82, plain, ((~ (e3 = op(e3, op(e3, e1))) & ~ (e2 = op(e2, op(e3, e1))) & ~ (e1 = op(e1, op(e3, e1))) & ~ (e0 = op(e0, op(e3, e1))) & (op(e1, e3) = op(e3, e1))) | ~ sP28), inference(usedef, [], [e82])).
fof(e82, plain, (sP28 <=> (~ (e3 = op(e3, op(e3, e1))) & ~ (e2 = op(e2, op(e3, e1))) & ~ (e1 = op(e1, op(e3, e1))) & ~ (e0 = op(e0, op(e3, e1))) & (op(e1, e3) = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP28])])).
fof(f83, plain, ((~ (e3 = op(e3, op(e3, e2))) & ~ (e2 = op(e2, op(e3, e2))) & ~ (e1 = op(e1, op(e3, e2))) & ~ (e0 = op(e0, op(e3, e2))) & (op(e2, e3) = op(e3, e2))) | ~ sP29), inference(usedef, [], [e83])).
fof(e83, plain, (sP29 <=> (~ (e3 = op(e3, op(e3, e2))) & ~ (e2 = op(e2, op(e3, e2))) & ~ (e1 = op(e1, op(e3, e2))) & ~ (e0 = op(e0, op(e3, e2))) & (op(e2, e3) = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP29])])).
fof(f5, plain, (((~ (e3 = op(e3, op(e3, e3))) & ~ (e2 = op(e2, op(e3, e3))) & ~ (e1 = op(e1, op(e3, e3))) & ~ (e0 = op(e0, op(e3, e3))) & (op(e3, e3) = op(e3, e3))) | (~ (e3 = op(e3, op(e3, e2))) & ~ (e2 = op(e2, op(e3, e2))) & ~ (e1 = op(e1, op(e3, e2))) & ~ (e0 = op(e0, op(e3, e2))) & (op(e2, e3) = op(e3, e2))) | (~ (e3 = op(e3, op(e3, e1))) & ~ (e2 = op(e2, op(e3, e1))) & ~ (e1 = op(e1, op(e3, e1))) & ~ (e0 = op(e0, op(e3, e1))) & (op(e1, e3) = op(e3, e1))) | (~ (e3 = op(e3, op(e3, e0))) & ~ (e2 = op(e2, op(e3, e0))) & ~ (e1 = op(e1, op(e3, e0))) & ~ (e0 = op(e0, op(e3, e0))) & (op(e0, e3) = op(e3, e0))) | (~ (e3 = op(e3, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e3))) & ~ (e1 = op(e1, op(e2, e3))) & ~ (e0 = op(e0, op(e2, e3))) & (op(e2, e3) = op(e3, e2))) | (~ (e3 = op(e3, op(e2, e2))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e1, op(e2, e2))) & ~ (e0 = op(e0, op(e2, e2))) & (op(e2, e2) = op(e2, e2))) | (~ (e3 = op(e3, op(e2, e1))) & ~ (e2 = op(e2, op(e2, e1))) & ~ (e1 = op(e1, op(e2, e1))) & ~ (e0 = op(e0, op(e2, e1))) & (op(e1, e2) = op(e2, e1))) | (~ (e3 = op(e3, op(e2, e0))) & ~ (e2 = op(e2, op(e2, e0))) & ~ (e1 = op(e1, op(e2, e0))) & ~ (e0 = op(e0, op(e2, e0))) & (op(e0, e2) = op(e2, e0))) | (~ (e3 = op(e3, op(e1, e3))) & ~ (e2 = op(e2, op(e1, e3))) & ~ (e1 = op(e1, op(e1, e3))) & ~ (e0 = op(e0, op(e1, e3))) & (op(e1, e3) = op(e3, e1))) | (~ (e3 = op(e3, op(e1, e2))) & ~ (e2 = op(e2, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e2))) & ~ (e0 = op(e0, op(e1, e2))) & (op(e1, e2) = op(e2, e1))) | (~ (e3 = op(e3, op(e1, e1))) & ~ (e2 = op(e2, op(e1, e1))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e0, op(e1, e1))) & (op(e1, e1) = op(e1, e1))) | (~ (e3 = op(e3, op(e1, e0))) & ~ (e2 = op(e2, op(e1, e0))) & ~ (e1 = op(e1, op(e1, e0))) & ~ (e0 = op(e0, op(e1, e0))) & (op(e0, e1) = op(e1, e0))) | (~ (e3 = op(e3, op(e0, e3))) & ~ (e2 = op(e2, op(e0, e3))) & ~ (e1 = op(e1, op(e0, e3))) & ~ (e0 = op(e0, op(e0, e3))) & (op(e0, e3) = op(e3, e0))) | (~ (e3 = op(e3, op(e0, e2))) & ~ (e2 = op(e2, op(e0, e2))) & ~ (e1 = op(e1, op(e0, e2))) & ~ (e0 = op(e0, op(e0, e2))) & (op(e0, e2) = op(e2, e0))) | (~ (e3 = op(e3, op(e0, e1))) & ~ (e2 = op(e2, op(e0, e1))) & ~ (e1 = op(e1, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e1))) & (op(e0, e1) = op(e1, e0))) | (~ (e3 = op(e3, op(e0, e0))) & ~ (e2 = op(e2, op(e0, e0))) & ~ (e1 = op(e1, op(e0, e0))) & ~ (e0 = op(e0, op(e0, e0))) & (op(e0, e0) = op(e0, e0)))) & ((~ (e3 = op(op(e3, e3), e3)) & (e3 = op(op(e3, e3), e3)) & ~ (e3 = op(e3, op(e3, e3)))) | (~ (e3 = op(op(e3, e2), e2)) & (e2 = op(op(e3, e2), e3)) & ~ (e2 = op(e3, op(e3, e2)))) | (~ (e3 = op(op(e3, e1), e1)) & (e1 = op(op(e3, e1), e3)) & ~ (e1 = op(e3, op(e3, e1)))) | (~ (e3 = op(op(e3, e0), e0)) & (e0 = op(op(e3, e0), e3)) & ~ (e0 = op(e3, op(e3, e0)))) | (~ (e2 = op(op(e2, e3), e3)) & (e3 = op(op(e2, e3), e2)) & ~ (e3 = op(e2, op(e2, e3)))) | (~ (e2 = op(op(e2, e2), e2)) & (e2 = op(op(e2, e2), e2)) & ~ (e2 = op(e2, op(e2, e2)))) | (~ (e2 = op(op(e2, e1), e1)) & (e1 = op(op(e2, e1), e2)) & ~ (e1 = op(e2, op(e2, e1)))) | (~ (e2 = op(op(e2, e0), e0)) & (e0 = op(op(e2, e0), e2)) & ~ (e0 = op(e2, op(e2, e0)))) | (~ (e1 = op(op(e1, e3), e3)) & (e3 = op(op(e1, e3), e1)) & ~ (e3 = op(e1, op(e1, e3)))) | (~ (e1 = op(op(e1, e2), e2)) & (e2 = op(op(e1, e2), e1)) & ~ (e2 = op(e1, op(e1, e2)))) | (~ (e1 = op(op(e1, e1), e1)) & (e1 = op(op(e1, e1), e1)) & ~ (e1 = op(e1, op(e1, e1)))) | (~ (e1 = op(op(e1, e0), e0)) & (e0 = op(op(e1, e0), e1)) & ~ (e0 = op(e1, op(e1, e0)))) | (~ (e0 = op(op(e0, e3), e3)) & (e3 = op(op(e0, e3), e0)) & ~ (e3 = op(e0, op(e0, e3)))) | (~ (e0 = op(op(e0, e2), e2)) & (e2 = op(op(e0, e2), e0)) & ~ (e2 = op(e0, op(e0, e2)))) | (~ (e0 = op(op(e0, e1), e1)) & (e1 = op(op(e0, e1), e0)) & ~ (e1 = op(e0, op(e0, e1)))) | (~ (e0 = op(op(e0, e0), e0)) & (e0 = op(op(e0, e0), e0)) & ~ (e0 = op(e0, op(e0, e0)))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG157+1.p', ax5)).
fof(f1355, plain, (spl30_198 | spl30_194 | spl30_190 | spl30_186 | spl30_182 | spl30_180 | spl30_176 | spl30_172 | spl30_168 | spl30_164 | spl30_162 | spl30_158 | spl30_154 | spl30_150 | spl30_146 | spl30_204), inference(avatar_split_clause, [], [f338, f1351, f1070, f1089, f1108, f1127, f1146, f1157, f1176, f1195, f1214, f1233, f1244, f1263, f1282, f1301, f1320])).
fof(f1351, plain, (spl30_204 <=> (e3 = op(op(e3, e3), e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_204])])).
fof(f338, plain, ((e3 = op(op(e3, e3), e3)) | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f84])).
fof(f1354, plain, (spl30_198 | spl30_194 | spl30_190 | spl30_186 | spl30_182 | spl30_180 | spl30_176 | spl30_172 | spl30_168 | spl30_164 | spl30_162 | spl30_158 | spl30_154 | spl30_150 | spl30_146 | ~ spl30_204), inference(avatar_split_clause, [], [f339, f1351, f1070, f1089, f1108, f1127, f1146, f1157, f1176, f1195, f1214, f1233, f1244, f1263, f1282, f1301, f1320])).
fof(f339, plain, (~ (e3 = op(op(e3, e3), e3)) | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f84])).
fof(f1349, plain, (spl30_141 | spl30_136 | spl30_131 | spl30_126 | spl30_120 | spl30_115 | spl30_110 | spl30_105 | spl30_99 | spl30_93 | spl30_88 | spl30_83 | spl30_77 | spl30_71 | spl30_65 | ~ spl30_203), inference(avatar_split_clause, [], [f341, f1346, f674, f703, f732, f761, f786, f810, f839, f868, f893, f918, f942, f971, f996, f1021, f1046])).
fof(f1046, plain, (spl30_141 <=> sP15), introduced(avatar_definition, [new_symbols(naming, [spl30_141])])).
fof(f1021, plain, (spl30_136 <=> sP16), introduced(avatar_definition, [new_symbols(naming, [spl30_136])])).
fof(f996, plain, (spl30_131 <=> sP17), introduced(avatar_definition, [new_symbols(naming, [spl30_131])])).
fof(f971, plain, (spl30_126 <=> sP18), introduced(avatar_definition, [new_symbols(naming, [spl30_126])])).
fof(f942, plain, (spl30_120 <=> sP19), introduced(avatar_definition, [new_symbols(naming, [spl30_120])])).
fof(f918, plain, (spl30_115 <=> sP20), introduced(avatar_definition, [new_symbols(naming, [spl30_115])])).
fof(f893, plain, (spl30_110 <=> sP21), introduced(avatar_definition, [new_symbols(naming, [spl30_110])])).
fof(f868, plain, (spl30_105 <=> sP22), introduced(avatar_definition, [new_symbols(naming, [spl30_105])])).
fof(f839, plain, (spl30_99 <=> sP23), introduced(avatar_definition, [new_symbols(naming, [spl30_99])])).
fof(f810, plain, (spl30_93 <=> sP24), introduced(avatar_definition, [new_symbols(naming, [spl30_93])])).
fof(f786, plain, (spl30_88 <=> sP25), introduced(avatar_definition, [new_symbols(naming, [spl30_88])])).
fof(f761, plain, (spl30_83 <=> sP26), introduced(avatar_definition, [new_symbols(naming, [spl30_83])])).
fof(f732, plain, (spl30_77 <=> sP27), introduced(avatar_definition, [new_symbols(naming, [spl30_77])])).
fof(f703, plain, (spl30_71 <=> sP28), introduced(avatar_definition, [new_symbols(naming, [spl30_71])])).
fof(f674, plain, (spl30_65 <=> sP29), introduced(avatar_definition, [new_symbols(naming, [spl30_65])])).
fof(f341, plain, (~ (e0 = op(e0, op(e3, e3))) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15), inference(cnf_transformation, [], [f84])).
fof(f1339, plain, (spl30_141 | spl30_136 | spl30_131 | spl30_126 | spl30_120 | spl30_115 | spl30_110 | spl30_105 | spl30_99 | spl30_93 | spl30_88 | spl30_83 | spl30_77 | spl30_71 | spl30_65 | ~ spl30_201), inference(avatar_split_clause, [], [f343, f1336, f674, f703, f732, f761, f786, f810, f839, f868, f893, f918, f942, f971, f996, f1021, f1046])).
fof(f343, plain, (~ (e2 = op(e2, op(e3, e3))) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15), inference(cnf_transformation, [], [f84])).
fof(f1334, plain, (spl30_141 | spl30_136 | spl30_131 | spl30_126 | spl30_120 | spl30_115 | spl30_110 | spl30_105 | spl30_99 | spl30_93 | spl30_88 | spl30_83 | spl30_77 | spl30_71 | spl30_65 | ~ spl30_200), inference(avatar_split_clause, [], [f344, f1331, f674, f703, f732, f761, f786, f810, f839, f868, f893, f918, f942, f971, f996, f1021, f1046])).
fof(f344, plain, (~ (e3 = op(e3, op(e3, e3))) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15), inference(cnf_transformation, [], [f84])).
fof(f1328, plain, (~ spl30_198 | spl30_199), inference(avatar_split_clause, [], [f335, f1324, f1320])).
fof(f1324, plain, (spl30_199 <=> (e0 = op(op(e0, e0), e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_199])])).
fof(f335, plain, ((e0 = op(op(e0, e0), e0)) | ~ sP0), inference(cnf_transformation, [], [f114])).
fof(f114, plain, ((~ (e0 = op(op(e0, e0), e0)) & (e0 = op(op(e0, e0), e0)) & ~ (e0 = op(e0, op(e0, e0)))) | ~ sP0), inference(nnf_transformation, [], [f54])).
fof(f1327, plain, (~ spl30_198 | ~ spl30_199), inference(avatar_split_clause, [], [f336, f1324, f1320])).
fof(f336, plain, (~ (e0 = op(op(e0, e0), e0)) | ~ sP0), inference(cnf_transformation, [], [f114])).
fof(f1318, plain, (~ spl30_194 | ~ spl30_197), inference(avatar_split_clause, [], [f331, f1315, f1301])).
fof(f331, plain, (~ (e1 = op(e0, op(e0, e1))) | ~ sP1), inference(cnf_transformation, [], [f113])).
fof(f113, plain, ((~ (e0 = op(op(e0, e1), e1)) & (e1 = op(op(e0, e1), e0)) & ~ (e1 = op(e0, op(e0, e1)))) | ~ sP1), inference(nnf_transformation, [], [f55])).
fof(f1313, plain, (~ spl30_194 | spl30_196), inference(avatar_split_clause, [], [f332, f1310, f1301])).
fof(f332, plain, ((e1 = op(op(e0, e1), e0)) | ~ sP1), inference(cnf_transformation, [], [f113])).
fof(f1308, plain, (~ spl30_194 | ~ spl30_195), inference(avatar_split_clause, [], [f333, f1305, f1301])).
fof(f333, plain, (~ (e0 = op(op(e0, e1), e1)) | ~ sP1), inference(cnf_transformation, [], [f113])).
fof(f1299, plain, (~ spl30_190 | ~ spl30_193), inference(avatar_split_clause, [], [f328, f1296, f1282])).
fof(f328, plain, (~ (e2 = op(e0, op(e0, e2))) | ~ sP2), inference(cnf_transformation, [], [f112])).
fof(f112, plain, ((~ (e0 = op(op(e0, e2), e2)) & (e2 = op(op(e0, e2), e0)) & ~ (e2 = op(e0, op(e0, e2)))) | ~ sP2), inference(nnf_transformation, [], [f56])).
fof(f1294, plain, (~ spl30_190 | spl30_192), inference(avatar_split_clause, [], [f329, f1291, f1282])).
fof(f329, plain, ((e2 = op(op(e0, e2), e0)) | ~ sP2), inference(cnf_transformation, [], [f112])).
fof(f1289, plain, (~ spl30_190 | ~ spl30_191), inference(avatar_split_clause, [], [f330, f1286, f1282])).
fof(f330, plain, (~ (e0 = op(op(e0, e2), e2)) | ~ sP2), inference(cnf_transformation, [], [f112])).
fof(f1280, plain, (~ spl30_186 | ~ spl30_189), inference(avatar_split_clause, [], [f325, f1277, f1263])).
fof(f325, plain, (~ (e3 = op(e0, op(e0, e3))) | ~ sP3), inference(cnf_transformation, [], [f111])).
fof(f111, plain, ((~ (e0 = op(op(e0, e3), e3)) & (e3 = op(op(e0, e3), e0)) & ~ (e3 = op(e0, op(e0, e3)))) | ~ sP3), inference(nnf_transformation, [], [f57])).
fof(f1275, plain, (~ spl30_186 | spl30_188), inference(avatar_split_clause, [], [f326, f1272, f1263])).
fof(f326, plain, ((e3 = op(op(e0, e3), e0)) | ~ sP3), inference(cnf_transformation, [], [f111])).
fof(f1270, plain, (~ spl30_186 | ~ spl30_187), inference(avatar_split_clause, [], [f327, f1267, f1263])).
fof(f327, plain, (~ (e0 = op(op(e0, e3), e3)) | ~ sP3), inference(cnf_transformation, [], [f111])).
fof(f1261, plain, (~ spl30_182 | ~ spl30_185), inference(avatar_split_clause, [], [f322, f1258, f1244])).
fof(f322, plain, (~ (e0 = op(e1, op(e1, e0))) | ~ sP4), inference(cnf_transformation, [], [f110])).
fof(f110, plain, ((~ (e1 = op(op(e1, e0), e0)) & (e0 = op(op(e1, e0), e1)) & ~ (e0 = op(e1, op(e1, e0)))) | ~ sP4), inference(nnf_transformation, [], [f58])).
fof(f1256, plain, (~ spl30_182 | spl30_184), inference(avatar_split_clause, [], [f323, f1253, f1244])).
fof(f323, plain, ((e0 = op(op(e1, e0), e1)) | ~ sP4), inference(cnf_transformation, [], [f110])).
fof(f1251, plain, (~ spl30_182 | ~ spl30_183), inference(avatar_split_clause, [], [f324, f1248, f1244])).
fof(f324, plain, (~ (e1 = op(op(e1, e0), e0)) | ~ sP4), inference(cnf_transformation, [], [f110])).
fof(f1241, plain, (~ spl30_180 | spl30_181), inference(avatar_split_clause, [], [f320, f1237, f1233])).
fof(f1237, plain, (spl30_181 <=> (e1 = op(op(e1, e1), e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_181])])).
fof(f320, plain, ((e1 = op(op(e1, e1), e1)) | ~ sP5), inference(cnf_transformation, [], [f109])).
fof(f109, plain, ((~ (e1 = op(op(e1, e1), e1)) & (e1 = op(op(e1, e1), e1)) & ~ (e1 = op(e1, op(e1, e1)))) | ~ sP5), inference(nnf_transformation, [], [f59])).
fof(f1240, plain, (~ spl30_180 | ~ spl30_181), inference(avatar_split_clause, [], [f321, f1237, f1233])).
fof(f321, plain, (~ (e1 = op(op(e1, e1), e1)) | ~ sP5), inference(cnf_transformation, [], [f109])).
fof(f1231, plain, (~ spl30_176 | ~ spl30_179), inference(avatar_split_clause, [], [f316, f1228, f1214])).
fof(f316, plain, (~ (e2 = op(e1, op(e1, e2))) | ~ sP6), inference(cnf_transformation, [], [f108])).
fof(f108, plain, ((~ (e1 = op(op(e1, e2), e2)) & (e2 = op(op(e1, e2), e1)) & ~ (e2 = op(e1, op(e1, e2)))) | ~ sP6), inference(nnf_transformation, [], [f60])).
fof(f1226, plain, (~ spl30_176 | spl30_178), inference(avatar_split_clause, [], [f317, f1223, f1214])).
fof(f317, plain, ((e2 = op(op(e1, e2), e1)) | ~ sP6), inference(cnf_transformation, [], [f108])).
fof(f1221, plain, (~ spl30_176 | ~ spl30_177), inference(avatar_split_clause, [], [f318, f1218, f1214])).
fof(f318, plain, (~ (e1 = op(op(e1, e2), e2)) | ~ sP6), inference(cnf_transformation, [], [f108])).
fof(f1212, plain, (~ spl30_172 | ~ spl30_175), inference(avatar_split_clause, [], [f313, f1209, f1195])).
fof(f313, plain, (~ (e3 = op(e1, op(e1, e3))) | ~ sP7), inference(cnf_transformation, [], [f107])).
fof(f107, plain, ((~ (e1 = op(op(e1, e3), e3)) & (e3 = op(op(e1, e3), e1)) & ~ (e3 = op(e1, op(e1, e3)))) | ~ sP7), inference(nnf_transformation, [], [f61])).
fof(f1207, plain, (~ spl30_172 | spl30_174), inference(avatar_split_clause, [], [f314, f1204, f1195])).
fof(f314, plain, ((e3 = op(op(e1, e3), e1)) | ~ sP7), inference(cnf_transformation, [], [f107])).
fof(f1202, plain, (~ spl30_172 | ~ spl30_173), inference(avatar_split_clause, [], [f315, f1199, f1195])).
fof(f315, plain, (~ (e1 = op(op(e1, e3), e3)) | ~ sP7), inference(cnf_transformation, [], [f107])).
fof(f1193, plain, (~ spl30_168 | ~ spl30_171), inference(avatar_split_clause, [], [f310, f1190, f1176])).
fof(f310, plain, (~ (e0 = op(e2, op(e2, e0))) | ~ sP8), inference(cnf_transformation, [], [f106])).
fof(f106, plain, ((~ (e2 = op(op(e2, e0), e0)) & (e0 = op(op(e2, e0), e2)) & ~ (e0 = op(e2, op(e2, e0)))) | ~ sP8), inference(nnf_transformation, [], [f62])).
fof(f1188, plain, (~ spl30_168 | spl30_170), inference(avatar_split_clause, [], [f311, f1185, f1176])).
fof(f311, plain, ((e0 = op(op(e2, e0), e2)) | ~ sP8), inference(cnf_transformation, [], [f106])).
fof(f1183, plain, (~ spl30_168 | ~ spl30_169), inference(avatar_split_clause, [], [f312, f1180, f1176])).
fof(f312, plain, (~ (e2 = op(op(e2, e0), e0)) | ~ sP8), inference(cnf_transformation, [], [f106])).
fof(f1174, plain, (~ spl30_164 | ~ spl30_167), inference(avatar_split_clause, [], [f307, f1171, f1157])).
fof(f307, plain, (~ (e1 = op(e2, op(e2, e1))) | ~ sP9), inference(cnf_transformation, [], [f105])).
fof(f105, plain, ((~ (e2 = op(op(e2, e1), e1)) & (e1 = op(op(e2, e1), e2)) & ~ (e1 = op(e2, op(e2, e1)))) | ~ sP9), inference(nnf_transformation, [], [f63])).
fof(f1169, plain, (~ spl30_164 | spl30_166), inference(avatar_split_clause, [], [f308, f1166, f1157])).
fof(f308, plain, ((e1 = op(op(e2, e1), e2)) | ~ sP9), inference(cnf_transformation, [], [f105])).
fof(f1164, plain, (~ spl30_164 | ~ spl30_165), inference(avatar_split_clause, [], [f309, f1161, f1157])).
fof(f309, plain, (~ (e2 = op(op(e2, e1), e1)) | ~ sP9), inference(cnf_transformation, [], [f105])).
fof(f1154, plain, (~ spl30_162 | spl30_163), inference(avatar_split_clause, [], [f305, f1150, f1146])).
fof(f1150, plain, (spl30_163 <=> (e2 = op(op(e2, e2), e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_163])])).
fof(f305, plain, ((e2 = op(op(e2, e2), e2)) | ~ sP10), inference(cnf_transformation, [], [f104])).
fof(f104, plain, ((~ (e2 = op(op(e2, e2), e2)) & (e2 = op(op(e2, e2), e2)) & ~ (e2 = op(e2, op(e2, e2)))) | ~ sP10), inference(nnf_transformation, [], [f64])).
fof(f1153, plain, (~ spl30_162 | ~ spl30_163), inference(avatar_split_clause, [], [f306, f1150, f1146])).
fof(f306, plain, (~ (e2 = op(op(e2, e2), e2)) | ~ sP10), inference(cnf_transformation, [], [f104])).
fof(f1144, plain, (~ spl30_158 | ~ spl30_161), inference(avatar_split_clause, [], [f301, f1141, f1127])).
fof(f301, plain, (~ (e3 = op(e2, op(e2, e3))) | ~ sP11), inference(cnf_transformation, [], [f103])).
fof(f103, plain, ((~ (e2 = op(op(e2, e3), e3)) & (e3 = op(op(e2, e3), e2)) & ~ (e3 = op(e2, op(e2, e3)))) | ~ sP11), inference(nnf_transformation, [], [f65])).
fof(f1139, plain, (~ spl30_158 | spl30_160), inference(avatar_split_clause, [], [f302, f1136, f1127])).
fof(f302, plain, ((e3 = op(op(e2, e3), e2)) | ~ sP11), inference(cnf_transformation, [], [f103])).
fof(f1134, plain, (~ spl30_158 | ~ spl30_159), inference(avatar_split_clause, [], [f303, f1131, f1127])).
fof(f303, plain, (~ (e2 = op(op(e2, e3), e3)) | ~ sP11), inference(cnf_transformation, [], [f103])).
fof(f1125, plain, (~ spl30_154 | ~ spl30_157), inference(avatar_split_clause, [], [f298, f1122, f1108])).
fof(f298, plain, (~ (e0 = op(e3, op(e3, e0))) | ~ sP12), inference(cnf_transformation, [], [f102])).
fof(f102, plain, ((~ (e3 = op(op(e3, e0), e0)) & (e0 = op(op(e3, e0), e3)) & ~ (e0 = op(e3, op(e3, e0)))) | ~ sP12), inference(nnf_transformation, [], [f66])).
fof(f1120, plain, (~ spl30_154 | spl30_156), inference(avatar_split_clause, [], [f299, f1117, f1108])).
fof(f299, plain, ((e0 = op(op(e3, e0), e3)) | ~ sP12), inference(cnf_transformation, [], [f102])).
fof(f1115, plain, (~ spl30_154 | ~ spl30_155), inference(avatar_split_clause, [], [f300, f1112, f1108])).
fof(f300, plain, (~ (e3 = op(op(e3, e0), e0)) | ~ sP12), inference(cnf_transformation, [], [f102])).
fof(f1106, plain, (~ spl30_150 | ~ spl30_153), inference(avatar_split_clause, [], [f295, f1103, f1089])).
fof(f295, plain, (~ (e1 = op(e3, op(e3, e1))) | ~ sP13), inference(cnf_transformation, [], [f101])).
fof(f101, plain, ((~ (e3 = op(op(e3, e1), e1)) & (e1 = op(op(e3, e1), e3)) & ~ (e1 = op(e3, op(e3, e1)))) | ~ sP13), inference(nnf_transformation, [], [f67])).
fof(f1101, plain, (~ spl30_150 | spl30_152), inference(avatar_split_clause, [], [f296, f1098, f1089])).
fof(f296, plain, ((e1 = op(op(e3, e1), e3)) | ~ sP13), inference(cnf_transformation, [], [f101])).
fof(f1096, plain, (~ spl30_150 | ~ spl30_151), inference(avatar_split_clause, [], [f297, f1093, f1089])).
fof(f297, plain, (~ (e3 = op(op(e3, e1), e1)) | ~ sP13), inference(cnf_transformation, [], [f101])).
fof(f1087, plain, (~ spl30_146 | ~ spl30_149), inference(avatar_split_clause, [], [f292, f1084, f1070])).
fof(f292, plain, (~ (e2 = op(e3, op(e3, e2))) | ~ sP14), inference(cnf_transformation, [], [f100])).
fof(f100, plain, ((~ (e3 = op(op(e3, e2), e2)) & (e2 = op(op(e3, e2), e3)) & ~ (e2 = op(e3, op(e3, e2)))) | ~ sP14), inference(nnf_transformation, [], [f68])).
fof(f1082, plain, (~ spl30_146 | spl30_148), inference(avatar_split_clause, [], [f293, f1079, f1070])).
fof(f293, plain, ((e2 = op(op(e3, e2), e3)) | ~ sP14), inference(cnf_transformation, [], [f100])).
fof(f1077, plain, (~ spl30_146 | ~ spl30_147), inference(avatar_split_clause, [], [f294, f1074, f1070])).
fof(f294, plain, (~ (e3 = op(op(e3, e2), e2)) | ~ sP14), inference(cnf_transformation, [], [f100])).
fof(f1068, plain, (~ spl30_141 | ~ spl30_145), inference(avatar_split_clause, [], [f288, f1065, f1046])).
fof(f288, plain, (~ (e0 = op(e0, op(e0, e0))) | ~ sP15), inference(cnf_transformation, [], [f99])).
fof(f99, plain, ((~ (e3 = op(e3, op(e0, e0))) & ~ (e2 = op(e2, op(e0, e0))) & ~ (e1 = op(e1, op(e0, e0))) & ~ (e0 = op(e0, op(e0, e0))) & (op(e0, e0) = op(e0, e0))) | ~ sP15), inference(nnf_transformation, [], [f69])).
fof(f1063, plain, (~ spl30_141 | ~ spl30_144), inference(avatar_split_clause, [], [f289, f1060, f1046])).
fof(f289, plain, (~ (e1 = op(e1, op(e0, e0))) | ~ sP15), inference(cnf_transformation, [], [f99])).
fof(f1058, plain, (~ spl30_141 | ~ spl30_143), inference(avatar_split_clause, [], [f290, f1055, f1046])).
fof(f290, plain, (~ (e2 = op(e2, op(e0, e0))) | ~ sP15), inference(cnf_transformation, [], [f99])).
fof(f1053, plain, (~ spl30_141 | ~ spl30_142), inference(avatar_split_clause, [], [f291, f1050, f1046])).
fof(f291, plain, (~ (e3 = op(e3, op(e0, e0))) | ~ sP15), inference(cnf_transformation, [], [f99])).
fof(f1044, plain, (~ spl30_136 | spl30_125), inference(avatar_split_clause, [], [f282, f966, f1021])).
fof(f282, plain, ((op(e0, e1) = op(e1, e0)) | ~ sP16), inference(cnf_transformation, [], [f98])).
fof(f98, plain, ((~ (e3 = op(e3, op(e0, e1))) & ~ (e2 = op(e2, op(e0, e1))) & ~ (e1 = op(e1, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e1))) & (op(e0, e1) = op(e1, e0))) | ~ sP16), inference(nnf_transformation, [], [f70])).
fof(f1038, plain, (~ spl30_136 | ~ spl30_139), inference(avatar_split_clause, [], [f284, f1035, f1021])).
fof(f284, plain, (~ (e1 = op(e1, op(e0, e1))) | ~ sP16), inference(cnf_transformation, [], [f98])).
fof(f1033, plain, (~ spl30_136 | ~ spl30_138), inference(avatar_split_clause, [], [f285, f1030, f1021])).
fof(f285, plain, (~ (e2 = op(e2, op(e0, e1))) | ~ sP16), inference(cnf_transformation, [], [f98])).
fof(f1028, plain, (~ spl30_136 | ~ spl30_137), inference(avatar_split_clause, [], [f286, f1025, f1021])).
fof(f286, plain, (~ (e3 = op(e3, op(e0, e1))) | ~ sP16), inference(cnf_transformation, [], [f98])).
fof(f1019, plain, (~ spl30_131 | spl30_104), inference(avatar_split_clause, [], [f277, f863, f996])).
fof(f277, plain, ((op(e0, e2) = op(e2, e0)) | ~ sP17), inference(cnf_transformation, [], [f97])).
fof(f97, plain, ((~ (e3 = op(e3, op(e0, e2))) & ~ (e2 = op(e2, op(e0, e2))) & ~ (e1 = op(e1, op(e0, e2))) & ~ (e0 = op(e0, op(e0, e2))) & (op(e0, e2) = op(e2, e0))) | ~ sP17), inference(nnf_transformation, [], [f71])).
fof(f1018, plain, (~ spl30_131 | ~ spl30_135), inference(avatar_split_clause, [], [f278, f1015, f996])).
fof(f278, plain, (~ (e0 = op(e0, op(e0, e2))) | ~ sP17), inference(cnf_transformation, [], [f97])).
fof(f1013, plain, (~ spl30_131 | ~ spl30_134), inference(avatar_split_clause, [], [f279, f1010, f996])).
fof(f279, plain, (~ (e1 = op(e1, op(e0, e2))) | ~ sP17), inference(cnf_transformation, [], [f97])).
fof(f1008, plain, (~ spl30_131 | ~ spl30_133), inference(avatar_split_clause, [], [f280, f1005, f996])).
fof(f280, plain, (~ (e2 = op(e2, op(e0, e2))) | ~ sP17), inference(cnf_transformation, [], [f97])).
fof(f1003, plain, (~ spl30_131 | ~ spl30_132), inference(avatar_split_clause, [], [f281, f1000, f996])).
fof(f281, plain, (~ (e3 = op(e3, op(e0, e2))) | ~ sP17), inference(cnf_transformation, [], [f97])).
fof(f994, plain, (~ spl30_126 | spl30_82), inference(avatar_split_clause, [], [f272, f756, f971])).
fof(f272, plain, ((op(e0, e3) = op(e3, e0)) | ~ sP18), inference(cnf_transformation, [], [f96])).
fof(f96, plain, ((~ (e3 = op(e3, op(e0, e3))) & ~ (e2 = op(e2, op(e0, e3))) & ~ (e1 = op(e1, op(e0, e3))) & ~ (e0 = op(e0, op(e0, e3))) & (op(e0, e3) = op(e3, e0))) | ~ sP18), inference(nnf_transformation, [], [f72])).
fof(f993, plain, (~ spl30_126 | ~ spl30_130), inference(avatar_split_clause, [], [f273, f990, f971])).
fof(f273, plain, (~ (e0 = op(e0, op(e0, e3))) | ~ sP18), inference(cnf_transformation, [], [f96])).
fof(f988, plain, (~ spl30_126 | ~ spl30_129), inference(avatar_split_clause, [], [f274, f985, f971])).
fof(f274, plain, (~ (e1 = op(e1, op(e0, e3))) | ~ sP18), inference(cnf_transformation, [], [f96])).
fof(f983, plain, (~ spl30_126 | ~ spl30_128), inference(avatar_split_clause, [], [f275, f980, f971])).
fof(f275, plain, (~ (e2 = op(e2, op(e0, e3))) | ~ sP18), inference(cnf_transformation, [], [f96])).
fof(f978, plain, (~ spl30_126 | ~ spl30_127), inference(avatar_split_clause, [], [f276, f975, f971])).
fof(f276, plain, (~ (e3 = op(e3, op(e0, e3))) | ~ sP18), inference(cnf_transformation, [], [f96])).
fof(f969, plain, (~ spl30_120 | spl30_125), inference(avatar_split_clause, [], [f267, f966, f942])).
fof(f267, plain, ((op(e0, e1) = op(e1, e0)) | ~ sP19), inference(cnf_transformation, [], [f95])).
fof(f95, plain, ((~ (e3 = op(e3, op(e1, e0))) & ~ (e2 = op(e2, op(e1, e0))) & ~ (e1 = op(e1, op(e1, e0))) & ~ (e0 = op(e0, op(e1, e0))) & (op(e0, e1) = op(e1, e0))) | ~ sP19), inference(nnf_transformation, [], [f73])).
fof(f959, plain, (~ spl30_120 | ~ spl30_123), inference(avatar_split_clause, [], [f269, f956, f942])).
fof(f269, plain, (~ (e1 = op(e1, op(e1, e0))) | ~ sP19), inference(cnf_transformation, [], [f95])).
fof(f954, plain, (~ spl30_120 | ~ spl30_122), inference(avatar_split_clause, [], [f270, f951, f942])).
fof(f270, plain, (~ (e2 = op(e2, op(e1, e0))) | ~ sP19), inference(cnf_transformation, [], [f95])).
fof(f949, plain, (~ spl30_120 | ~ spl30_121), inference(avatar_split_clause, [], [f271, f946, f942])).
fof(f271, plain, (~ (e3 = op(e3, op(e1, e0))) | ~ sP19), inference(cnf_transformation, [], [f95])).
fof(f940, plain, (~ spl30_115 | ~ spl30_119), inference(avatar_split_clause, [], [f263, f937, f918])).
fof(f263, plain, (~ (e0 = op(e0, op(e1, e1))) | ~ sP20), inference(cnf_transformation, [], [f94])).
fof(f94, plain, ((~ (e3 = op(e3, op(e1, e1))) & ~ (e2 = op(e2, op(e1, e1))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e0, op(e1, e1))) & (op(e1, e1) = op(e1, e1))) | ~ sP20), inference(nnf_transformation, [], [f74])).
fof(f935, plain, (~ spl30_115 | ~ spl30_118), inference(avatar_split_clause, [], [f264, f932, f918])).
fof(f264, plain, (~ (e1 = op(e1, op(e1, e1))) | ~ sP20), inference(cnf_transformation, [], [f94])).
fof(f930, plain, (~ spl30_115 | ~ spl30_117), inference(avatar_split_clause, [], [f265, f927, f918])).
fof(f265, plain, (~ (e2 = op(e2, op(e1, e1))) | ~ sP20), inference(cnf_transformation, [], [f94])).
fof(f925, plain, (~ spl30_115 | ~ spl30_116), inference(avatar_split_clause, [], [f266, f922, f918])).
fof(f266, plain, (~ (e3 = op(e3, op(e1, e1))) | ~ sP20), inference(cnf_transformation, [], [f94])).
fof(f916, plain, (~ spl30_110 | spl30_98), inference(avatar_split_clause, [], [f257, f834, f893])).
fof(f257, plain, ((op(e1, e2) = op(e2, e1)) | ~ sP21), inference(cnf_transformation, [], [f93])).
fof(f93, plain, ((~ (e3 = op(e3, op(e1, e2))) & ~ (e2 = op(e2, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e2))) & ~ (e0 = op(e0, op(e1, e2))) & (op(e1, e2) = op(e2, e1))) | ~ sP21), inference(nnf_transformation, [], [f75])).
fof(f915, plain, (~ spl30_110 | ~ spl30_114), inference(avatar_split_clause, [], [f258, f912, f893])).
fof(f258, plain, (~ (e0 = op(e0, op(e1, e2))) | ~ sP21), inference(cnf_transformation, [], [f93])).
fof(f910, plain, (~ spl30_110 | ~ spl30_113), inference(avatar_split_clause, [], [f259, f907, f893])).
fof(f259, plain, (~ (e1 = op(e1, op(e1, e2))) | ~ sP21), inference(cnf_transformation, [], [f93])).
fof(f905, plain, (~ spl30_110 | ~ spl30_112), inference(avatar_split_clause, [], [f260, f902, f893])).
fof(f260, plain, (~ (e2 = op(e2, op(e1, e2))) | ~ sP21), inference(cnf_transformation, [], [f93])).
fof(f900, plain, (~ spl30_110 | ~ spl30_111), inference(avatar_split_clause, [], [f261, f897, f893])).
fof(f261, plain, (~ (e3 = op(e3, op(e1, e2))) | ~ sP21), inference(cnf_transformation, [], [f93])).
fof(f891, plain, (~ spl30_105 | spl30_76), inference(avatar_split_clause, [], [f252, f727, f868])).
fof(f252, plain, ((op(e1, e3) = op(e3, e1)) | ~ sP22), inference(cnf_transformation, [], [f92])).
fof(f92, plain, ((~ (e3 = op(e3, op(e1, e3))) & ~ (e2 = op(e2, op(e1, e3))) & ~ (e1 = op(e1, op(e1, e3))) & ~ (e0 = op(e0, op(e1, e3))) & (op(e1, e3) = op(e3, e1))) | ~ sP22), inference(nnf_transformation, [], [f76])).
fof(f890, plain, (~ spl30_105 | ~ spl30_109), inference(avatar_split_clause, [], [f253, f887, f868])).
fof(f253, plain, (~ (e0 = op(e0, op(e1, e3))) | ~ sP22), inference(cnf_transformation, [], [f92])).
fof(f885, plain, (~ spl30_105 | ~ spl30_108), inference(avatar_split_clause, [], [f254, f882, f868])).
fof(f254, plain, (~ (e1 = op(e1, op(e1, e3))) | ~ sP22), inference(cnf_transformation, [], [f92])).
fof(f880, plain, (~ spl30_105 | ~ spl30_107), inference(avatar_split_clause, [], [f255, f877, f868])).
fof(f255, plain, (~ (e2 = op(e2, op(e1, e3))) | ~ sP22), inference(cnf_transformation, [], [f92])).
fof(f866, plain, (~ spl30_99 | spl30_104), inference(avatar_split_clause, [], [f247, f863, f839])).
fof(f247, plain, ((op(e0, e2) = op(e2, e0)) | ~ sP23), inference(cnf_transformation, [], [f91])).
fof(f91, plain, ((~ (e3 = op(e3, op(e2, e0))) & ~ (e2 = op(e2, op(e2, e0))) & ~ (e1 = op(e1, op(e2, e0))) & ~ (e0 = op(e0, op(e2, e0))) & (op(e0, e2) = op(e2, e0))) | ~ sP23), inference(nnf_transformation, [], [f77])).
fof(f856, plain, (~ spl30_99 | ~ spl30_102), inference(avatar_split_clause, [], [f249, f853, f839])).
fof(f249, plain, (~ (e1 = op(e1, op(e2, e0))) | ~ sP23), inference(cnf_transformation, [], [f91])).
fof(f851, plain, (~ spl30_99 | ~ spl30_101), inference(avatar_split_clause, [], [f250, f848, f839])).
fof(f250, plain, (~ (e2 = op(e2, op(e2, e0))) | ~ sP23), inference(cnf_transformation, [], [f91])).
fof(f846, plain, (~ spl30_99 | ~ spl30_100), inference(avatar_split_clause, [], [f251, f843, f839])).
fof(f251, plain, (~ (e3 = op(e3, op(e2, e0))) | ~ sP23), inference(cnf_transformation, [], [f91])).
fof(f837, plain, (~ spl30_93 | spl30_98), inference(avatar_split_clause, [], [f242, f834, f810])).
fof(f242, plain, ((op(e1, e2) = op(e2, e1)) | ~ sP24), inference(cnf_transformation, [], [f90])).
fof(f90, plain, ((~ (e3 = op(e3, op(e2, e1))) & ~ (e2 = op(e2, op(e2, e1))) & ~ (e1 = op(e1, op(e2, e1))) & ~ (e0 = op(e0, op(e2, e1))) & (op(e1, e2) = op(e2, e1))) | ~ sP24), inference(nnf_transformation, [], [f78])).
fof(f832, plain, (~ spl30_93 | ~ spl30_97), inference(avatar_split_clause, [], [f243, f829, f810])).
fof(f243, plain, (~ (e0 = op(e0, op(e2, e1))) | ~ sP24), inference(cnf_transformation, [], [f90])).
fof(f822, plain, (~ spl30_93 | ~ spl30_95), inference(avatar_split_clause, [], [f245, f819, f810])).
fof(f245, plain, (~ (e2 = op(e2, op(e2, e1))) | ~ sP24), inference(cnf_transformation, [], [f90])).
fof(f817, plain, (~ spl30_93 | ~ spl30_94), inference(avatar_split_clause, [], [f246, f814, f810])).
fof(f246, plain, (~ (e3 = op(e3, op(e2, e1))) | ~ sP24), inference(cnf_transformation, [], [f90])).
fof(f808, plain, (~ spl30_88 | ~ spl30_92), inference(avatar_split_clause, [], [f238, f805, f786])).
fof(f238, plain, (~ (e0 = op(e0, op(e2, e2))) | ~ sP25), inference(cnf_transformation, [], [f89])).
fof(f89, plain, ((~ (e3 = op(e3, op(e2, e2))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e1, op(e2, e2))) & ~ (e0 = op(e0, op(e2, e2))) & (op(e2, e2) = op(e2, e2))) | ~ sP25), inference(nnf_transformation, [], [f79])).
fof(f803, plain, (~ spl30_88 | ~ spl30_91), inference(avatar_split_clause, [], [f239, f800, f786])).
fof(f239, plain, (~ (e1 = op(e1, op(e2, e2))) | ~ sP25), inference(cnf_transformation, [], [f89])).
fof(f798, plain, (~ spl30_88 | ~ spl30_90), inference(avatar_split_clause, [], [f240, f795, f786])).
fof(f240, plain, (~ (e2 = op(e2, op(e2, e2))) | ~ sP25), inference(cnf_transformation, [], [f89])).
fof(f793, plain, (~ spl30_88 | ~ spl30_89), inference(avatar_split_clause, [], [f241, f790, f786])).
fof(f241, plain, (~ (e3 = op(e3, op(e2, e2))) | ~ sP25), inference(cnf_transformation, [], [f89])).
fof(f784, plain, (~ spl30_83 | spl30_70), inference(avatar_split_clause, [], [f232, f698, f761])).
fof(f232, plain, ((op(e2, e3) = op(e3, e2)) | ~ sP26), inference(cnf_transformation, [], [f88])).
fof(f88, plain, ((~ (e3 = op(e3, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e3))) & ~ (e1 = op(e1, op(e2, e3))) & ~ (e0 = op(e0, op(e2, e3))) & (op(e2, e3) = op(e3, e2))) | ~ sP26), inference(nnf_transformation, [], [f80])).
fof(f783, plain, (~ spl30_83 | ~ spl30_87), inference(avatar_split_clause, [], [f233, f780, f761])).
fof(f233, plain, (~ (e0 = op(e0, op(e2, e3))) | ~ sP26), inference(cnf_transformation, [], [f88])).
fof(f778, plain, (~ spl30_83 | ~ spl30_86), inference(avatar_split_clause, [], [f234, f775, f761])).
fof(f234, plain, (~ (e1 = op(e1, op(e2, e3))) | ~ sP26), inference(cnf_transformation, [], [f88])).
fof(f773, plain, (~ spl30_83 | ~ spl30_85), inference(avatar_split_clause, [], [f235, f770, f761])).
fof(f235, plain, (~ (e2 = op(e2, op(e2, e3))) | ~ sP26), inference(cnf_transformation, [], [f88])).
fof(f768, plain, (~ spl30_83 | ~ spl30_84), inference(avatar_split_clause, [], [f236, f765, f761])).
fof(f236, plain, (~ (e3 = op(e3, op(e2, e3))) | ~ sP26), inference(cnf_transformation, [], [f88])).
fof(f759, plain, (~ spl30_77 | spl30_82), inference(avatar_split_clause, [], [f227, f756, f732])).
fof(f227, plain, ((op(e0, e3) = op(e3, e0)) | ~ sP27), inference(cnf_transformation, [], [f87])).
fof(f87, plain, ((~ (e3 = op(e3, op(e3, e0))) & ~ (e2 = op(e2, op(e3, e0))) & ~ (e1 = op(e1, op(e3, e0))) & ~ (e0 = op(e0, op(e3, e0))) & (op(e0, e3) = op(e3, e0))) | ~ sP27), inference(nnf_transformation, [], [f81])).
fof(f754, plain, (~ spl30_77 | ~ spl30_81), inference(avatar_split_clause, [], [f228, f751, f732])).
fof(f228, plain, (~ (e0 = op(e0, op(e3, e0))) | ~ sP27), inference(cnf_transformation, [], [f87])).
fof(f749, plain, (~ spl30_77 | ~ spl30_80), inference(avatar_split_clause, [], [f229, f746, f732])).
fof(f229, plain, (~ (e1 = op(e1, op(e3, e0))) | ~ sP27), inference(cnf_transformation, [], [f87])).
fof(f739, plain, (~ spl30_77 | ~ spl30_78), inference(avatar_split_clause, [], [f231, f736, f732])).
fof(f231, plain, (~ (e3 = op(e3, op(e3, e0))) | ~ sP27), inference(cnf_transformation, [], [f87])).
fof(f730, plain, (~ spl30_71 | spl30_76), inference(avatar_split_clause, [], [f222, f727, f703])).
fof(f222, plain, ((op(e1, e3) = op(e3, e1)) | ~ sP28), inference(cnf_transformation, [], [f86])).
fof(f86, plain, ((~ (e3 = op(e3, op(e3, e1))) & ~ (e2 = op(e2, op(e3, e1))) & ~ (e1 = op(e1, op(e3, e1))) & ~ (e0 = op(e0, op(e3, e1))) & (op(e1, e3) = op(e3, e1))) | ~ sP28), inference(nnf_transformation, [], [f82])).
fof(f725, plain, (~ spl30_71 | ~ spl30_75), inference(avatar_split_clause, [], [f223, f722, f703])).
fof(f223, plain, (~ (e0 = op(e0, op(e3, e1))) | ~ sP28), inference(cnf_transformation, [], [f86])).
fof(f720, plain, (~ spl30_71 | ~ spl30_74), inference(avatar_split_clause, [], [f224, f717, f703])).
fof(f224, plain, (~ (e1 = op(e1, op(e3, e1))) | ~ sP28), inference(cnf_transformation, [], [f86])).
fof(f715, plain, (~ spl30_71 | ~ spl30_73), inference(avatar_split_clause, [], [f225, f712, f703])).
fof(f225, plain, (~ (e2 = op(e2, op(e3, e1))) | ~ sP28), inference(cnf_transformation, [], [f86])).
fof(f701, plain, (~ spl30_65 | spl30_70), inference(avatar_split_clause, [], [f217, f698, f674])).
fof(f217, plain, ((op(e2, e3) = op(e3, e2)) | ~ sP29), inference(cnf_transformation, [], [f85])).
fof(f85, plain, ((~ (e3 = op(e3, op(e3, e2))) & ~ (e2 = op(e2, op(e3, e2))) & ~ (e1 = op(e1, op(e3, e2))) & ~ (e0 = op(e0, op(e3, e2))) & (op(e2, e3) = op(e3, e2))) | ~ sP29), inference(nnf_transformation, [], [f83])).
fof(f696, plain, (~ spl30_65 | ~ spl30_69), inference(avatar_split_clause, [], [f218, f693, f674])).
fof(f218, plain, (~ (e0 = op(e0, op(e3, e2))) | ~ sP29), inference(cnf_transformation, [], [f85])).
fof(f691, plain, (~ spl30_65 | ~ spl30_68), inference(avatar_split_clause, [], [f219, f688, f674])).
fof(f219, plain, (~ (e1 = op(e1, op(e3, e2))) | ~ sP29), inference(cnf_transformation, [], [f85])).
fof(f686, plain, (~ spl30_65 | ~ spl30_67), inference(avatar_split_clause, [], [f220, f683, f674])).
fof(f220, plain, (~ (e2 = op(e2, op(e3, e2))) | ~ sP29), inference(cnf_transformation, [], [f85])).
fof(f681, plain, (~ spl30_65 | ~ spl30_66), inference(avatar_split_clause, [], [f221, f678, f674])).
fof(f221, plain, (~ (e3 = op(e3, op(e3, e2))) | ~ sP29), inference(cnf_transformation, [], [f85])).
fof(f672, plain, (spl30_61 | spl30_57 | spl30_53 | spl30_49), inference(avatar_split_clause, [], [f131, f574, f591, f608, f625])).
fof(f131, plain, ((e0 = op(e0, e3)) | (e0 = op(e0, e2)) | (e0 = op(e0, e1)) | (e0 = op(e0, e0))), inference(cnf_transformation, [], [f2])).
fof(f2, plain, (((e3 = op(e3, e3)) | (e3 = op(e2, e3)) | (e3 = op(e1, e3)) | (e3 = op(e0, e3))) & ((e3 = op(e3, e3)) | (e3 = op(e3, e2)) | (e3 = op(e3, e1)) | (e3 = op(e3, e0))) & ((e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(e0, e3))) & ((e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, e0))) & ((e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(e0, e3))) & ((e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, e0))) & ((e0 = op(e3, e3)) | (e0 = op(e2, e3)) | (e0 = op(e1, e3)) | (e0 = op(e0, e3))) & ((e0 = op(e3, e3)) | (e0 = op(e3, e2)) | (e0 = op(e3, e1)) | (e0 = op(e3, e0))) & ((e3 = op(e3, e2)) | (e3 = op(e2, e2)) | (e3 = op(e1, e2)) | (e3 = op(e0, e2))) & ((e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, e0))) & ((e2 = op(e3, e2)) | (e2 = op(e2, e2)) | (e2 = op(e1, e2)) | (e2 = op(e0, e2))) & ((e2 = op(e2, e3)) | (e2 = op(e2, e2)) | (e2 = op(e2, e1)) | (e2 = op(e2, e0))) & ((e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(e0, e2))) & ((e1 = op(e2, e3)) | (e1 = op(e2, e2)) | (e1 = op(e2, e1)) | (e1 = op(e2, e0))) & ((e0 = op(e3, e2)) | (e0 = op(e2, e2)) | (e0 = op(e1, e2)) | (e0 = op(e0, e2))) & ((e0 = op(e2, e3)) | (e0 = op(e2, e2)) | (e0 = op(e2, e1)) | (e0 = op(e2, e0))) & ((e3 = op(e3, e1)) | (e3 = op(e2, e1)) | (e3 = op(e1, e1)) | (e3 = op(e0, e1))) & ((e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, e0))) & ((e2 = op(e3, e1)) | (e2 = op(e2, e1)) | (e2 = op(e1, e1)) | (e2 = op(e0, e1))) & ((e2 = op(e1, e3)) | (e2 = op(e1, e2)) | (e2 = op(e1, e1)) | (e2 = op(e1, e0))) & ((e1 = op(e3, e1)) | (e1 = op(e2, e1)) | (e1 = op(e1, e1)) | (e1 = op(e0, e1))) & ((e1 = op(e1, e3)) | (e1 = op(e1, e2)) | (e1 = op(e1, e1)) | (e1 = op(e1, e0))) & ((e0 = op(e3, e1)) | (e0 = op(e2, e1)) | (e0 = op(e1, e1)) | (e0 = op(e0, e1))) & ((e0 = op(e1, e3)) | (e0 = op(e1, e2)) | (e0 = op(e1, e1)) | (e0 = op(e1, e0))) & ((e3 = op(e3, e0)) | (e3 = op(e2, e0)) | (e3 = op(e1, e0)) | (op(e0, e0) = e3)) & ((e3 = op(e0, e3)) | (e3 = op(e0, e2)) | (e3 = op(e0, e1)) | (op(e0, e0) = e3)) & ((e2 = op(e3, e0)) | (e2 = op(e2, e0)) | (e2 = op(e1, e0)) | (op(e0, e0) = e2)) & ((e2 = op(e0, e3)) | (e2 = op(e0, e2)) | (e2 = op(e0, e1)) | (op(e0, e0) = e2)) & ((e1 = op(e3, e0)) | (e1 = op(e2, e0)) | (e1 = op(e1, e0)) | (op(e0, e0) = e1)) & ((e1 = op(e0, e3)) | (e1 = op(e0, e2)) | (e1 = op(e0, e1)) | (op(e0, e0) = e1)) & ((e0 = op(e3, e0)) | (e0 = op(e2, e0)) | (e0 = op(e1, e0)) | (e0 = op(e0, e0))) & ((e0 = op(e0, e3)) | (e0 = op(e0, e2)) | (e0 = op(e0, e1)) | (e0 = op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG157+1.p', ax2)).
fof(f671, plain, (spl30_61 | spl30_45 | spl30_29 | spl30_13), inference(avatar_split_clause, [], [f132, f421, f489, f557, f625])).
fof(f132, plain, ((e0 = op(e3, e0)) | (e0 = op(e2, e0)) | (e0 = op(e1, e0)) | (e0 = op(e0, e0))), inference(cnf_transformation, [], [f2])).
fof(f670, plain, (spl30_62 | spl30_58 | spl30_54 | spl30_50), inference(avatar_split_clause, [], [f133, f578, f595, f612, f629])).
fof(f133, plain, ((e1 = op(e0, e3)) | (e1 = op(e0, e2)) | (e1 = op(e0, e1)) | (op(e0, e0) = e1)), inference(cnf_transformation, [], [f2])).
fof(f669, plain, (spl30_62 | spl30_46 | spl30_30 | spl30_14), inference(avatar_split_clause, [], [f134, f425, f493, f561, f629])).
fof(f134, plain, ((e1 = op(e3, e0)) | (e1 = op(e2, e0)) | (e1 = op(e1, e0)) | (op(e0, e0) = e1)), inference(cnf_transformation, [], [f2])).
fof(f668, plain, (spl30_63 | spl30_59 | spl30_55 | spl30_51), inference(avatar_split_clause, [], [f135, f582, f599, f616, f633])).
fof(f135, plain, ((e2 = op(e0, e3)) | (e2 = op(e0, e2)) | (e2 = op(e0, e1)) | (op(e0, e0) = e2)), inference(cnf_transformation, [], [f2])).
fof(f667, plain, (spl30_63 | spl30_47 | spl30_31 | spl30_15), inference(avatar_split_clause, [], [f136, f429, f497, f565, f633])).
fof(f136, plain, ((e2 = op(e3, e0)) | (e2 = op(e2, e0)) | (e2 = op(e1, e0)) | (op(e0, e0) = e2)), inference(cnf_transformation, [], [f2])).
fof(f666, plain, (spl30_64 | spl30_60 | spl30_56 | spl30_52), inference(avatar_split_clause, [], [f137, f586, f603, f620, f637])).
fof(f137, plain, ((e3 = op(e0, e3)) | (e3 = op(e0, e2)) | (e3 = op(e0, e1)) | (op(e0, e0) = e3)), inference(cnf_transformation, [], [f2])).
fof(f665, plain, (spl30_64 | spl30_48 | spl30_32 | spl30_16), inference(avatar_split_clause, [], [f138, f433, f501, f569, f637])).
fof(f138, plain, ((e3 = op(e3, e0)) | (e3 = op(e2, e0)) | (e3 = op(e1, e0)) | (op(e0, e0) = e3)), inference(cnf_transformation, [], [f2])).
fof(f664, plain, (spl30_45 | spl30_41 | spl30_37 | spl30_33), inference(avatar_split_clause, [], [f139, f506, f523, f540, f557])).
fof(f139, plain, ((e0 = op(e1, e3)) | (e0 = op(e1, e2)) | (e0 = op(e1, e1)) | (e0 = op(e1, e0))), inference(cnf_transformation, [], [f2])).
fof(f663, plain, (spl30_57 | spl30_41 | spl30_25 | spl30_9), inference(avatar_split_clause, [], [f140, f404, f472, f540, f608])).
fof(f140, plain, ((e0 = op(e3, e1)) | (e0 = op(e2, e1)) | (e0 = op(e1, e1)) | (e0 = op(e0, e1))), inference(cnf_transformation, [], [f2])).
fof(f662, plain, (spl30_46 | spl30_42 | spl30_38 | spl30_34), inference(avatar_split_clause, [], [f141, f510, f527, f544, f561])).
fof(f141, plain, ((e1 = op(e1, e3)) | (e1 = op(e1, e2)) | (e1 = op(e1, e1)) | (e1 = op(e1, e0))), inference(cnf_transformation, [], [f2])).
fof(f661, plain, (spl30_58 | spl30_42 | spl30_26 | spl30_10), inference(avatar_split_clause, [], [f142, f408, f476, f544, f612])).
fof(f142, plain, ((e1 = op(e3, e1)) | (e1 = op(e2, e1)) | (e1 = op(e1, e1)) | (e1 = op(e0, e1))), inference(cnf_transformation, [], [f2])).
fof(f660, plain, (spl30_47 | spl30_43 | spl30_39 | spl30_35), inference(avatar_split_clause, [], [f143, f514, f531, f548, f565])).
fof(f143, plain, ((e2 = op(e1, e3)) | (e2 = op(e1, e2)) | (e2 = op(e1, e1)) | (e2 = op(e1, e0))), inference(cnf_transformation, [], [f2])).
fof(f659, plain, (spl30_59 | spl30_43 | spl30_27 | spl30_11), inference(avatar_split_clause, [], [f144, f412, f480, f548, f616])).
fof(f144, plain, ((e2 = op(e3, e1)) | (e2 = op(e2, e1)) | (e2 = op(e1, e1)) | (e2 = op(e0, e1))), inference(cnf_transformation, [], [f2])).
fof(f658, plain, (spl30_48 | spl30_44 | spl30_40 | spl30_36), inference(avatar_split_clause, [], [f145, f518, f535, f552, f569])).
fof(f145, plain, ((e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, e0))), inference(cnf_transformation, [], [f2])).
fof(f657, plain, (spl30_60 | spl30_44 | spl30_28 | spl30_12), inference(avatar_split_clause, [], [f146, f416, f484, f552, f620])).
fof(f146, plain, ((e3 = op(e3, e1)) | (e3 = op(e2, e1)) | (e3 = op(e1, e1)) | (e3 = op(e0, e1))), inference(cnf_transformation, [], [f2])).
fof(f656, plain, (spl30_29 | spl30_25 | spl30_21 | spl30_17), inference(avatar_split_clause, [], [f147, f438, f455, f472, f489])).
fof(f147, plain, ((e0 = op(e2, e3)) | (e0 = op(e2, e2)) | (e0 = op(e2, e1)) | (e0 = op(e2, e0))), inference(cnf_transformation, [], [f2])).
fof(f655, plain, (spl30_53 | spl30_37 | spl30_21 | spl30_5), inference(avatar_split_clause, [], [f148, f387, f455, f523, f591])).
fof(f148, plain, ((e0 = op(e3, e2)) | (e0 = op(e2, e2)) | (e0 = op(e1, e2)) | (e0 = op(e0, e2))), inference(cnf_transformation, [], [f2])).
fof(f654, plain, (spl30_30 | spl30_26 | spl30_22 | spl30_18), inference(avatar_split_clause, [], [f149, f442, f459, f476, f493])).
fof(f149, plain, ((e1 = op(e2, e3)) | (e1 = op(e2, e2)) | (e1 = op(e2, e1)) | (e1 = op(e2, e0))), inference(cnf_transformation, [], [f2])).
fof(f653, plain, (spl30_54 | spl30_38 | spl30_22 | spl30_6), inference(avatar_split_clause, [], [f150, f391, f459, f527, f595])).
fof(f150, plain, ((e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(e0, e2))), inference(cnf_transformation, [], [f2])).
fof(f652, plain, (spl30_31 | spl30_27 | spl30_23 | spl30_19), inference(avatar_split_clause, [], [f151, f446, f463, f480, f497])).
fof(f151, plain, ((e2 = op(e2, e3)) | (e2 = op(e2, e2)) | (e2 = op(e2, e1)) | (e2 = op(e2, e0))), inference(cnf_transformation, [], [f2])).
fof(f651, plain, (spl30_55 | spl30_39 | spl30_23 | spl30_7), inference(avatar_split_clause, [], [f152, f395, f463, f531, f599])).
fof(f152, plain, ((e2 = op(e3, e2)) | (e2 = op(e2, e2)) | (e2 = op(e1, e2)) | (e2 = op(e0, e2))), inference(cnf_transformation, [], [f2])).
fof(f650, plain, (spl30_32 | spl30_28 | spl30_24 | spl30_20), inference(avatar_split_clause, [], [f153, f450, f467, f484, f501])).
fof(f153, plain, ((e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, e0))), inference(cnf_transformation, [], [f2])).
fof(f649, plain, (spl30_56 | spl30_40 | spl30_24 | spl30_8), inference(avatar_split_clause, [], [f154, f399, f467, f535, f603])).
fof(f154, plain, ((e3 = op(e3, e2)) | (e3 = op(e2, e2)) | (e3 = op(e1, e2)) | (e3 = op(e0, e2))), inference(cnf_transformation, [], [f2])).
fof(f648, plain, (spl30_13 | spl30_9 | spl30_5 | spl30_1), inference(avatar_split_clause, [], [f155, f370, f387, f404, f421])).
fof(f155, plain, ((e0 = op(e3, e3)) | (e0 = op(e3, e2)) | (e0 = op(e3, e1)) | (e0 = op(e3, e0))), inference(cnf_transformation, [], [f2])).
fof(f647, plain, (spl30_49 | spl30_33 | spl30_17 | spl30_1), inference(avatar_split_clause, [], [f156, f370, f438, f506, f574])).
fof(f156, plain, ((e0 = op(e3, e3)) | (e0 = op(e2, e3)) | (e0 = op(e1, e3)) | (e0 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f646, plain, (spl30_14 | spl30_10 | spl30_6 | spl30_2), inference(avatar_split_clause, [], [f157, f374, f391, f408, f425])).
fof(f157, plain, ((e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, e0))), inference(cnf_transformation, [], [f2])).
fof(f645, plain, (spl30_50 | spl30_34 | spl30_18 | spl30_2), inference(avatar_split_clause, [], [f158, f374, f442, f510, f578])).
fof(f158, plain, ((e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f644, plain, (spl30_15 | spl30_11 | spl30_7 | spl30_3), inference(avatar_split_clause, [], [f159, f378, f395, f412, f429])).
fof(f159, plain, ((e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, e0))), inference(cnf_transformation, [], [f2])).
fof(f643, plain, (spl30_51 | spl30_35 | spl30_19 | spl30_3), inference(avatar_split_clause, [], [f160, f378, f446, f514, f582])).
fof(f160, plain, ((e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f642, plain, (spl30_16 | spl30_12 | spl30_8 | spl30_4), inference(avatar_split_clause, [], [f161, f382, f399, f416, f433])).
fof(f161, plain, ((e3 = op(e3, e3)) | (e3 = op(e3, e2)) | (e3 = op(e3, e1)) | (e3 = op(e3, e0))), inference(cnf_transformation, [], [f2])).
fof(f641, plain, (spl30_52 | spl30_36 | spl30_20 | spl30_4), inference(avatar_split_clause, [], [f162, f382, f450, f518, f586])).
fof(f162, plain, ((e3 = op(e3, e3)) | (e3 = op(e2, e3)) | (e3 = op(e1, e3)) | (e3 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f640, plain, (spl30_61 | spl30_62 | spl30_63 | spl30_64), inference(avatar_split_clause, [], [f115, f637, f633, f629, f625])).
fof(f115, plain, ((op(e0, e0) = e3) | (op(e0, e0) = e2) | (op(e0, e0) = e1) | (e0 = op(e0, e0))), inference(cnf_transformation, [], [f1])).
fof(f1, plain, (((e3 = op(e3, e3)) | (e2 = op(e3, e3)) | (e1 = op(e3, e3)) | (e0 = op(e3, e3))) & ((e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (e0 = op(e3, e2))) & ((e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (e0 = op(e3, e1))) & ((e3 = op(e3, e0)) | (e2 = op(e3, e0)) | (e1 = op(e3, e0)) | (e0 = op(e3, e0))) & ((e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))) & ((e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (e0 = op(e2, e2))) & ((e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))) & ((e3 = op(e2, e0)) | (e2 = op(e2, e0)) | (e1 = op(e2, e0)) | (e0 = op(e2, e0))) & ((e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (e0 = op(e1, e3))) & ((e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (e0 = op(e1, e2))) & ((e3 = op(e1, e1)) | (e2 = op(e1, e1)) | (e1 = op(e1, e1)) | (e0 = op(e1, e1))) & ((e3 = op(e1, e0)) | (e2 = op(e1, e0)) | (e1 = op(e1, e0)) | (e0 = op(e1, e0))) & ((e3 = op(e0, e3)) | (e2 = op(e0, e3)) | (e1 = op(e0, e3)) | (e0 = op(e0, e3))) & ((e3 = op(e0, e2)) | (e2 = op(e0, e2)) | (e1 = op(e0, e2)) | (e0 = op(e0, e2))) & ((e3 = op(e0, e1)) | (e2 = op(e0, e1)) | (e1 = op(e0, e1)) | (e0 = op(e0, e1))) & ((op(e0, e0) = e3) | (op(e0, e0) = e2) | (op(e0, e0) = e1) | (e0 = op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG157+1.p', ax1)).
fof(f623, plain, (spl30_57 | spl30_58 | spl30_59 | spl30_60), inference(avatar_split_clause, [], [f116, f620, f616, f612, f608])).
fof(f116, plain, ((e3 = op(e0, e1)) | (e2 = op(e0, e1)) | (e1 = op(e0, e1)) | (e0 = op(e0, e1))), inference(cnf_transformation, [], [f1])).
fof(f606, plain, (spl30_53 | spl30_54 | spl30_55 | spl30_56), inference(avatar_split_clause, [], [f117, f603, f599, f595, f591])).
fof(f117, plain, ((e3 = op(e0, e2)) | (e2 = op(e0, e2)) | (e1 = op(e0, e2)) | (e0 = op(e0, e2))), inference(cnf_transformation, [], [f1])).
fof(f589, plain, (spl30_49 | spl30_50 | spl30_51 | spl30_52), inference(avatar_split_clause, [], [f118, f586, f582, f578, f574])).
fof(f118, plain, ((e3 = op(e0, e3)) | (e2 = op(e0, e3)) | (e1 = op(e0, e3)) | (e0 = op(e0, e3))), inference(cnf_transformation, [], [f1])).
fof(f572, plain, (spl30_45 | spl30_46 | spl30_47 | spl30_48), inference(avatar_split_clause, [], [f119, f569, f565, f561, f557])).
fof(f119, plain, ((e3 = op(e1, e0)) | (e2 = op(e1, e0)) | (e1 = op(e1, e0)) | (e0 = op(e1, e0))), inference(cnf_transformation, [], [f1])).
fof(f555, plain, (spl30_41 | spl30_42 | spl30_43 | spl30_44), inference(avatar_split_clause, [], [f120, f552, f548, f544, f540])).
fof(f120, plain, ((e3 = op(e1, e1)) | (e2 = op(e1, e1)) | (e1 = op(e1, e1)) | (e0 = op(e1, e1))), inference(cnf_transformation, [], [f1])).
fof(f538, plain, (spl30_37 | spl30_38 | spl30_39 | spl30_40), inference(avatar_split_clause, [], [f121, f535, f531, f527, f523])).
fof(f121, plain, ((e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (e0 = op(e1, e2))), inference(cnf_transformation, [], [f1])).
fof(f521, plain, (spl30_33 | spl30_34 | spl30_35 | spl30_36), inference(avatar_split_clause, [], [f122, f518, f514, f510, f506])).
fof(f122, plain, ((e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (e0 = op(e1, e3))), inference(cnf_transformation, [], [f1])).
fof(f504, plain, (spl30_29 | spl30_30 | spl30_31 | spl30_32), inference(avatar_split_clause, [], [f123, f501, f497, f493, f489])).
fof(f123, plain, ((e3 = op(e2, e0)) | (e2 = op(e2, e0)) | (e1 = op(e2, e0)) | (e0 = op(e2, e0))), inference(cnf_transformation, [], [f1])).
fof(f487, plain, (spl30_25 | spl30_26 | spl30_27 | spl30_28), inference(avatar_split_clause, [], [f124, f484, f480, f476, f472])).
fof(f124, plain, ((e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))), inference(cnf_transformation, [], [f1])).
fof(f470, plain, (spl30_21 | spl30_22 | spl30_23 | spl30_24), inference(avatar_split_clause, [], [f125, f467, f463, f459, f455])).
fof(f125, plain, ((e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (e0 = op(e2, e2))), inference(cnf_transformation, [], [f1])).
fof(f453, plain, (spl30_17 | spl30_18 | spl30_19 | spl30_20), inference(avatar_split_clause, [], [f126, f450, f446, f442, f438])).
fof(f126, plain, ((e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))), inference(cnf_transformation, [], [f1])).
fof(f436, plain, (spl30_13 | spl30_14 | spl30_15 | spl30_16), inference(avatar_split_clause, [], [f127, f433, f429, f425, f421])).
fof(f127, plain, ((e3 = op(e3, e0)) | (e2 = op(e3, e0)) | (e1 = op(e3, e0)) | (e0 = op(e3, e0))), inference(cnf_transformation, [], [f1])).
fof(f419, plain, (spl30_9 | spl30_10 | spl30_11 | spl30_12), inference(avatar_split_clause, [], [f128, f416, f412, f408, f404])).
fof(f128, plain, ((e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (e0 = op(e3, e1))), inference(cnf_transformation, [], [f1])).
fof(f402, plain, (spl30_5 | spl30_6 | spl30_7 | spl30_8), inference(avatar_split_clause, [], [f129, f399, f395, f391, f387])).
fof(f129, plain, ((e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (e0 = op(e3, e2))), inference(cnf_transformation, [], [f1])).
fof(f385, plain, (spl30_1 | spl30_2 | spl30_3 | spl30_4), inference(avatar_split_clause, [], [f130, f382, f378, f374, f370])).
fof(f130, plain, ((e3 = op(e3, e3)) | (e2 = op(e3, e3)) | (e1 = op(e3, e3)) | (e0 = op(e3, e3))), inference(cnf_transformation, [], [f1])).