fof(f10816, plain, $false, inference(avatar_sat_refutation, [], [f746, f763, f780, f797, f814, f831, f848, f865, f882, f899, f916, f933, f950, f1002, f1003, f1004, f1005, f1006, f1007, f1009, f1010, f1012, f1013, f1015, f1016, f1017, f1019, f1020, f1021, f1022, f1023, f1024, f1026, f1050, f1067, f1084, f1101, f1118, f1135, f1152, f1169, f1186, f1203, f1220, f1237, f1254, f1306, f1307, f1308, f1309, f1310, f1311, f1313, f1314, f1315, f1316, f1317, f1319, f1320, f1321, f1323, f1325, f1326, f1327, f1328, f1329, f1330, f1366, f1380, f1395, f1414, f1424, f1449, f1463, f1492, f1502, f1531, f1545, f1556, f1575, f1580, f1595, f1600, f1634, f1653, f1659, f1678, f1684, f1708, f1723, f1742, f1747, f1752, f1766, f1771, f1780, f1785, f1790, f1799, f1804, f1809, f1818, f1819, f1829, f1839, f1853, f1858, f1867, f1872, f1877, f1886, f1891, f1896, f1905, f1906, f1907, f1916, f1921, f1926, f1935, f1959, f1973, f1978, f1992, f1993, f2004, f2014, f2019, f2020, f2021, f2050, f2064, f2079, f2098, f2108, f2133, f2147, f2176, f2186, f2210, f2215, f2229, f2240, f2259, f2265, f2279, f2284, f2318, f2337, f2343, f2368, f2393, f2407, f2426, f2431, f2436, f2445, f2450, f2455, f2464, f2469, f2474, f2483, f2488, f2493, f2502, f2503, f2513, f2518, f2523, f2532, f2537, f2542, f2551, f2556, f2561, f2570, f2575, f2580, f2589, f2590, f2600, f2605, f2610, f2619, f2643, f2657, f2662, f2676, f2677, f2683, f2688, f2703, f2704, f2705, f2706, f2708, f2709, f2711, f2949, f2967, f2985, f3413, f3444, f3446, f3483, f3490, f3501, f3526, f3536, f3544, f3547, f3565, f3566, f3567, f3572, f3577, f3588, f3589, f3590, f3591, f3592, f3593, f3594, f3600, f3601, f3602, f3603, f3604, f3612, f3628, f3645, f3680, f3683, f3692, f3695, f3702, f3709, f3732, f3789, f3791, f3793, f3807, f3808, f3815, f3818, f3820, f3834, f3844, f3849, f3855, f3856, f3857, f3888, f3907, f3940, f3950, f3952, f3956, f3958, f3961, f3971, f3980, f3988, f4008, f4014, f4028, f4039, f4042, f4057, f4066, f4068, f4095, f4112, f4115, f4123, f4126, f4145, f4155, f4177, f4185, f4199, f4203, f4206, f4244, f4245, f4249, f4250, f4281, f4284, f4328, f4331, f4341, f4349, f4358, f4380, f4382, f4393, f4411, f4449, f4501, f4503, f4524, f4530, f4542, f4551, f4564, f4566, f4581, f4614, f4629, f4637, f4650, f4698, f4709, f4713, f4721, f4728, f4736, f4747, f4768, f4772, f4774, f4839, f4899, f4902, f4923, f4933, f4949, f4967, f4978, f4992, f4993, f5007, f5009, f5011, f5022, f5025, f5032, f5035, f5040, f5050, f5080, f5088, f5094, f5101, f5110, f5111, f5127, f5130, f5174, f5195, f5201, f5207, f5216, f5219, f5224, f5239, f5248, f5283, f5287, f5295, f5304, f5330, f5381, f5397, f5413, f5495, f5511, f5533, f5542, f5573, f5576, f5631, f5635, f5659, f5710, f5726, f5796, f5809, f5810, f5826, f5845, f5852, f5856, f5953, f5966, f6059, f6076, f6081, f6101, f6171, f6179, f6182, f6188, f6204, f6275, f6302, f6318, f6332, f6343, f6347, f6396, f6462, f6491, f6492, f6494, f6495, f6503, f6504, f6528, f6532, f6558, f6582, f6713, f6771, f6799, f6804, f6845, f6874, f6891, f6914, f6915, f6917, f6963, f6970, f6971, f7013, f7050, f7051, f7080, f7092, f7112, f7127, f7133, f7172, f7271, f7300, f7318, f7329, f7334, f7429, f7447, f7458, f7459, f7474, f7476, f7515, f7528, f7578, f7592, f7609, f7730, f7772, f7892, f7893, f7929, f7931, f7991, f8008, f8032, f8058, f8132, f8150, f8200, f8236, f8237, f8252, f8269, f8293, f8300, f8305, f8312, f8331, f8360, f8405, f8444, f8456, f8488, f8502, f8503, f8515, f8516, f8536, f8537, f8577, f8578, f8612, f8725, f8762, f8763, f8790, f8801, f8802, f8816, f8830, f8836, f8883, f8886, f8917, f8927, f8948, f8952, f8966, f8977, f8992, f8997, f9000, f9009, f9029, f9056, f9067, f9079, f9104, f9105, f9109, f9130, f9208, f9214, f9230, f9231, f9245, f9266, f9267, f9288, f9304, f9317, f9354, f9396, f9405, f9414, f9427, f9440, f9445, f9478, f9502, f9517, f9604, f9673, f9717, f9766, f9814, f9911, f9957, f9996, f9997, f10022, f10025, f10060, f10081, f10093, f10096, f10122, f10235, f10278, f10294, f10302, f10304, f10307, f10383, f10384, f10436, f10449, f10455, f10457, f10461, f10487, f10488, f10499, f10500, f10593, f10604, f10617, f10659, f10660, f10701, f10747, f10757, f10774, f10775, f10797])).
fof(f10797, plain, (spl72_92 | ~ spl72_99 | ~ spl72_378), inference(avatar_contradiction_clause, [], [f10796])).
fof(f10796, plain, ($false | (spl72_92 | ~ spl72_99 | ~ spl72_378)), inference(subsumption_resolution, [], [f10789, f1150])).
fof(f1150, plain, (~ (e23 = op2(e22, e21)) | spl72_92), inference(avatar_component_clause, [], [f1149])).
fof(f1149, plain, (spl72_92 <=> (e23 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl72_92])])).
fof(f10789, plain, ((e23 = op2(e22, e21)) | (~ spl72_99 | ~ spl72_378)), inference(backward_demodulation, [], [f2555, f1181])).
fof(f1181, plain, ((e22 = op2(e21, e23)) | ~ spl72_99), inference(avatar_component_clause, [], [f1179])).
fof(f1179, plain, (spl72_99 <=> (e22 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl72_99])])).
fof(f2555, plain, ((e23 = op2(op2(e21, e23), e21)) | ~ spl72_378), inference(avatar_component_clause, [], [f2553])).
fof(f2553, plain, (spl72_378 <=> (e23 = op2(op2(e21, e23), e21))), introduced(avatar_definition, [new_symbols(naming, [spl72_378])])).
fof(f10775, plain, (~ spl72_100 | ~ spl72_112), inference(avatar_split_clause, [], [f10768, f1234, f1183])).
fof(f1183, plain, (spl72_100 <=> (e23 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl72_100])])).
fof(f1234, plain, (spl72_112 <=> (e23 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl72_112])])).
fof(f10768, plain, (~ (e23 = op2(e21, e23)) | ~ spl72_112), inference(backward_demodulation, [], [f344, f1236])).
fof(f1236, plain, ((e23 = op2(e21, e20)) | ~ spl72_112), inference(avatar_component_clause, [], [f1234])).
fof(f344, plain, ~ (op2(e21, e20) = op2(e21, e23)), inference(cnf_transformation, [], [f6])).
fof(f6, plain, (~ (op2(e23, e22) = op2(e23, e23)) & ~ (op2(e23, e21) = op2(e23, e23)) & ~ (op2(e23, e21) = op2(e23, e22)) & ~ (op2(e23, e20) = op2(e23, e23)) & ~ (op2(e23, e20) = op2(e23, e22)) & ~ (op2(e23, e20) = op2(e23, e21)) & ~ (op2(e22, e22) = op2(e22, e23)) & ~ (op2(e22, e21) = op2(e22, e23)) & ~ (op2(e22, e21) = op2(e22, e22)) & ~ (op2(e22, e20) = op2(e22, e23)) & ~ (op2(e22, e20) = op2(e22, e22)) & ~ (op2(e22, e20) = op2(e22, e21)) & ~ (op2(e21, e22) = op2(e21, e23)) & ~ (op2(e21, e21) = op2(e21, e23)) & ~ (op2(e21, e21) = op2(e21, e22)) & ~ (op2(e21, e20) = op2(e21, e23)) & ~ (op2(e21, e20) = op2(e21, e22)) & ~ (op2(e21, e20) = op2(e21, e21)) & ~ (op2(e20, e22) = op2(e20, e23)) & ~ (op2(e20, e21) = op2(e20, e23)) & ~ (op2(e20, e21) = op2(e20, e22)) & ~ (op2(e20, e20) = op2(e20, e23)) & ~ (op2(e20, e20) = op2(e20, e22)) & ~ (op2(e20, e20) = op2(e20, e21)) & ~ (op2(e22, e23) = op2(e23, e23)) & ~ (op2(e21, e23) = op2(e23, e23)) & ~ (op2(e21, e23) = op2(e22, e23)) & ~ (op2(e20, e23) = op2(e23, e23)) & ~ (op2(e20, e23) = op2(e22, e23)) & ~ (op2(e20, e23) = op2(e21, e23)) & ~ (op2(e22, e22) = op2(e23, e22)) & ~ (op2(e21, e22) = op2(e23, e22)) & ~ (op2(e21, e22) = op2(e22, e22)) & ~ (op2(e20, e22) = op2(e23, e22)) & ~ (op2(e20, e22) = op2(e22, e22)) & ~ (op2(e20, e22) = op2(e21, e22)) & ~ (op2(e22, e21) = op2(e23, e21)) & ~ (op2(e21, e21) = op2(e23, e21)) & ~ (op2(e21, e21) = op2(e22, e21)) & ~ (op2(e20, e21) = op2(e23, e21)) & ~ (op2(e20, e21) = op2(e22, e21)) & ~ (op2(e20, e21) = op2(e21, e21)) & ~ (op2(e22, e20) = op2(e23, e20)) & ~ (op2(e21, e20) = op2(e23, e20)) & ~ (op2(e21, e20) = op2(e22, e20)) & ~ (op2(e20, e20) = op2(e23, e20)) & ~ (op2(e20, e20) = op2(e22, e20)) & ~ (op2(e20, e20) = op2(e21, e20))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG104+1.p', ax6)).
fof(f10774, plain, (~ spl72_80 | ~ spl72_112), inference(avatar_split_clause, [], [f10767, f1234, f1098])).
fof(f1098, plain, (spl72_80 <=> (e23 = op2(e23, e20))), introduced(avatar_definition, [new_symbols(naming, [spl72_80])])).
fof(f10767, plain, (~ (e23 = op2(e23, e20)) | ~ spl72_112), inference(backward_demodulation, [], [f316, f1236])).
fof(f316, plain, ~ (op2(e21, e20) = op2(e23, e20)), inference(cnf_transformation, [], [f6])).
fof(f10757, plain, (spl72_112 | ~ spl72_124 | ~ spl72_329), inference(avatar_split_clause, [], [f10756, f2315, f1285, f1234])).
fof(f1285, plain, (spl72_124 <=> (e23 = op2(e20, e21))), introduced(avatar_definition, [new_symbols(naming, [spl72_124])])).
fof(f2315, plain, (spl72_329 <=> (op2(e20, e21) = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl72_329])])).
fof(f10756, plain, ((e23 = op2(e21, e20)) | (~ spl72_124 | ~ spl72_329)), inference(forward_demodulation, [], [f2317, f1287])).
fof(f1287, plain, ((e23 = op2(e20, e21)) | ~ spl72_124), inference(avatar_component_clause, [], [f1285])).
fof(f2317, plain, ((op2(e20, e21) = op2(e21, e20)) | ~ spl72_329), inference(avatar_component_clause, [], [f2315])).
fof(f10747, plain, (spl72_78 | ~ spl72_124 | ~ spl72_400), inference(avatar_split_clause, [], [f10746, f2659, f1285, f1090])).
fof(f1090, plain, (spl72_78 <=> (e21 = op2(e23, e20))), introduced(avatar_definition, [new_symbols(naming, [spl72_78])])).
fof(f2659, plain, (spl72_400 <=> (e21 = op2(op2(e20, e21), e20))), introduced(avatar_definition, [new_symbols(naming, [spl72_400])])).
fof(f10746, plain, ((e21 = op2(e23, e20)) | (~ spl72_124 | ~ spl72_400)), inference(forward_demodulation, [], [f2661, f1287])).
fof(f2661, plain, ((e21 = op2(op2(e20, e21), e20)) | ~ spl72_400), inference(avatar_component_clause, [], [f2659])).
fof(f10701, plain, (~ spl72_90 | ~ spl72_439), inference(avatar_split_clause, [], [f10691, f2905, f1141])).
fof(f1141, plain, (spl72_90 <=> (e21 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl72_90])])).
fof(f2905, plain, (spl72_439 <=> (e21 = h2(e12))), introduced(avatar_definition, [new_symbols(naming, [spl72_439])])).
fof(f10691, plain, (~ (e21 = op2(e22, e21)) | ~ spl72_439), inference(backward_demodulation, [], [f2728, f2906])).
fof(f2906, plain, ((e21 = h2(e12)) | ~ spl72_439), inference(avatar_component_clause, [], [f2905])).
fof(f2728, plain, ~ (op2(e22, e21) = h2(e12)), inference(backward_demodulation, [], [f321, f656])).
fof(f656, plain, (op2(e21, e21) = h2(e12)), inference(cnf_transformation, [], [f15])).
fof(f15, plain, ((h2(e13) = op2(e21, op2(e21, op2(e21, e21)))) & (op2(e21, e21) = h2(e12)) & (op2(e21, op2(e21, e21)) = h2(e11)) & (e21 = h2(e10))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG104+1.p', ax15)).
fof(f321, plain, ~ (op2(e21, e21) = op2(e22, e21)), inference(cnf_transformation, [], [f6])).
fof(f10660, plain, (~ spl72_80 | ~ spl72_471), inference(avatar_split_clause, [], [f10649, f3076, f1098])).
fof(f3076, plain, (spl72_471 <=> (e23 = h4(e12))), introduced(avatar_definition, [new_symbols(naming, [spl72_471])])).
fof(f10649, plain, (~ (e23 = op2(e23, e20)) | ~ spl72_471), inference(backward_demodulation, [], [f2748, f3077])).
fof(f3077, plain, ((e23 = h4(e12)) | ~ spl72_471), inference(avatar_component_clause, [], [f3076])).
fof(f2748, plain, ~ (op2(e23, e20) = h4(e12)), inference(backward_demodulation, [], [f356, f664])).
fof(f664, plain, (op2(e23, e23) = h4(e12)), inference(cnf_transformation, [], [f17])).
fof(f17, plain, ((h4(e13) = op2(e23, op2(e23, op2(e23, e23)))) & (op2(e23, e23) = h4(e12)) & (op2(e23, op2(e23, e23)) = h4(e11)) & (e23 = h4(e10))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG104+1.p', ax17)).
fof(f356, plain, ~ (op2(e23, e20) = op2(e23, e23)), inference(cnf_transformation, [], [f6])).
fof(f10659, plain, (~ spl72_100 | ~ spl72_471), inference(avatar_split_clause, [], [f10648, f3076, f1183])).
fof(f10648, plain, (~ (e23 = op2(e21, e23)) | ~ spl72_471), inference(backward_demodulation, [], [f2746, f3077])).
fof(f2746, plain, ~ (op2(e21, e23) = h4(e12)), inference(backward_demodulation, [], [f334, f664])).
fof(f334, plain, ~ (op2(e21, e23) = op2(e23, e23)), inference(cnf_transformation, [], [f6])).
fof(f10617, plain, (~ spl72_83 | ~ spl72_19 | ~ spl72_447 | spl72_512 | ~ spl72_515), inference(avatar_split_clause, [], [f10616, f3362, f3350, f2946, f807, f1111])).
fof(f1111, plain, (spl72_83 <=> (e22 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl72_83])])).
fof(f807, plain, (spl72_19 <=> (e12 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl72_19])])).
fof(f2946, plain, (spl72_447 <=> (e22 = h1(e12))), introduced(avatar_definition, [new_symbols(naming, [spl72_447])])).
fof(f3350, plain, (spl72_512 <=> (h1(op1(e12, e13)) = op2(h1(e12), h1(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl72_512])])).
fof(f3362, plain, (spl72_515 <=> (e23 = h1(e13))), introduced(avatar_definition, [new_symbols(naming, [spl72_515])])).
fof(f10616, plain, (~ (e22 = op2(e22, e23)) | (~ spl72_19 | ~ spl72_447 | spl72_512 | ~ spl72_515)), inference(forward_demodulation, [], [f10615, f2947])).
fof(f2947, plain, ((e22 = h1(e12)) | ~ spl72_447), inference(avatar_component_clause, [], [f2946])).
fof(f10615, plain, (~ (op2(e22, e23) = h1(e12)) | (~ spl72_19 | ~ spl72_447 | spl72_512 | ~ spl72_515)), inference(forward_demodulation, [], [f10614, f809])).
fof(f809, plain, ((e12 = op1(e12, e13)) | ~ spl72_19), inference(avatar_component_clause, [], [f807])).
fof(f10614, plain, (~ (op2(e22, e23) = h1(op1(e12, e13))) | (~ spl72_447 | spl72_512 | ~ spl72_515)), inference(forward_demodulation, [], [f10613, f2947])).
fof(f10613, plain, (~ (h1(op1(e12, e13)) = op2(h1(e12), e23)) | (spl72_512 | ~ spl72_515)), inference(forward_demodulation, [], [f3352, f3363])).
fof(f3363, plain, ((e23 = h1(e13)) | ~ spl72_515), inference(avatar_component_clause, [], [f3362])).
fof(f3352, plain, (~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | spl72_512), inference(avatar_component_clause, [], [f3350])).
fof(f10604, plain, (~ spl72_90 | ~ spl72_26 | ~ spl72_447 | spl72_518), inference(avatar_split_clause, [], [f10603, f3374, f2946, f837, f1141])).
fof(f837, plain, (spl72_26 <=> (e11 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl72_26])])).
fof(f3374, plain, (spl72_518 <=> (h1(op1(e12, e11)) = op2(h1(e12), e21))), introduced(avatar_definition, [new_symbols(naming, [spl72_518])])).
fof(f10603, plain, (~ (e21 = op2(e22, e21)) | (~ spl72_26 | ~ spl72_447 | spl72_518)), inference(forward_demodulation, [], [f10602, f2727])).
fof(f2727, plain, (e21 = h1(e11)), inference(forward_demodulation, [], [f2726, f2724])).
fof(f2724, plain, (e21 = op2(e20, h1(e12))), inference(backward_demodulation, [], [f647, f652])).
fof(f652, plain, (op2(e20, e20) = h1(e12)), inference(cnf_transformation, [], [f14])).
fof(f14, plain, ((op2(e20, op2(e20, op2(e20, e20))) = h1(e13)) & (op2(e20, e20) = h1(e12)) & (op2(e20, op2(e20, e20)) = h1(e11)) & (e20 = h1(e10))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG104+1.p', ax14)).
fof(f647, plain, (e21 = op2(e20, op2(e20, e20))), inference(cnf_transformation, [], [f13])).
fof(f13, plain, ((e23 = op2(e20, op2(e20, op2(e20, e20)))) & (op2(e20, e20) = e22) & (e21 = op2(e20, op2(e20, e20)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG104+1.p', ax13)).
fof(f2726, plain, (h1(e11) = op2(e20, h1(e12))), inference(forward_demodulation, [], [f651, f652])).
fof(f651, plain, (op2(e20, op2(e20, e20)) = h1(e11)), inference(cnf_transformation, [], [f14])).
fof(f10602, plain, (~ (op2(e22, e21) = h1(e11)) | (~ spl72_26 | ~ spl72_447 | spl72_518)), inference(forward_demodulation, [], [f10601, f839])).
fof(f839, plain, ((e11 = op1(e12, e11)) | ~ spl72_26), inference(avatar_component_clause, [], [f837])).
fof(f10601, plain, (~ (op2(e22, e21) = h1(op1(e12, e11))) | (~ spl72_447 | spl72_518)), inference(forward_demodulation, [], [f3376, f2947])).
fof(f3376, plain, (~ (h1(op1(e12, e11)) = op2(h1(e12), e21)) | spl72_518), inference(avatar_component_clause, [], [f3374])).
fof(f10593, plain, (spl72_83 | ~ spl72_88 | ~ spl72_294), inference(avatar_split_clause, [], [f10592, f2144, f1132, f1111])).
fof(f1132, plain, (spl72_88 <=> (e23 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl72_88])])).
fof(f2144, plain, (spl72_294 <=> (e22 = op2(e22, op2(e22, e22)))), introduced(avatar_definition, [new_symbols(naming, [spl72_294])])).
fof(f10592, plain, ((e22 = op2(e22, e23)) | (~ spl72_88 | ~ spl72_294)), inference(forward_demodulation, [], [f2145, f1134])).
fof(f1134, plain, ((e23 = op2(e22, e22)) | ~ spl72_88), inference(avatar_component_clause, [], [f1132])).
fof(f2145, plain, ((e22 = op2(e22, op2(e22, e22))) | ~ spl72_294), inference(avatar_component_clause, [], [f2144])).
fof(f10500, plain, (~ spl72_82 | ~ spl72_90), inference(avatar_split_clause, [], [f10491, f1141, f1107])).
fof(f1107, plain, (spl72_82 <=> (e21 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl72_82])])).
fof(f10491, plain, (~ (e21 = op2(e22, e23)) | ~ spl72_90), inference(backward_demodulation, [], [f352, f1143])).
fof(f1143, plain, ((e21 = op2(e22, e21)) | ~ spl72_90), inference(avatar_component_clause, [], [f1141])).
fof(f352, plain, ~ (op2(e22, e21) = op2(e22, e23)), inference(cnf_transformation, [], [f6])).
fof(f10499, plain, (~ spl72_74 | ~ spl72_90), inference(avatar_split_clause, [], [f10490, f1141, f1073])).
fof(f1073, plain, (spl72_74 <=> (e21 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl72_74])])).
fof(f10490, plain, (~ (e21 = op2(e23, e21)) | ~ spl72_90), inference(backward_demodulation, [], [f323, f1143])).
fof(f323, plain, ~ (op2(e22, e21) = op2(e23, e21)), inference(cnf_transformation, [], [f6])).
fof(f10488, plain, (~ spl72_89 | ~ spl72_93), inference(avatar_split_clause, [], [f10476, f1154, f1137])).
fof(f1137, plain, (spl72_89 <=> (e20 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl72_89])])).
fof(f1154, plain, (spl72_93 <=> (e20 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl72_93])])).
fof(f10476, plain, (~ (e20 = op2(e22, e21)) | ~ spl72_93), inference(backward_demodulation, [], [f348, f1156])).
fof(f1156, plain, ((e20 = op2(e22, e20)) | ~ spl72_93), inference(avatar_component_clause, [], [f1154])).
fof(f348, plain, ~ (op2(e22, e20) = op2(e22, e21)), inference(cnf_transformation, [], [f6])).
fof(f10487, plain, (~ spl72_77 | ~ spl72_93), inference(avatar_split_clause, [], [f10475, f1154, f1086])).
fof(f1086, plain, (spl72_77 <=> (e20 = op2(e23, e20))), introduced(avatar_definition, [new_symbols(naming, [spl72_77])])).
fof(f10475, plain, (~ (e20 = op2(e23, e20)) | ~ spl72_93), inference(backward_demodulation, [], [f317, f1156])).
fof(f317, plain, ~ (op2(e22, e20) = op2(e23, e20)), inference(cnf_transformation, [], [f6])).
fof(f10461, plain, (~ spl72_89 | ~ spl72_101 | spl72_302), inference(avatar_split_clause, [], [f10460, f2183, f1188, f1137])).
fof(f1188, plain, (spl72_101 <=> (e20 = op2(e21, e22))), introduced(avatar_definition, [new_symbols(naming, [spl72_101])])).
fof(f2183, plain, (spl72_302 <=> (op2(e21, e22) = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl72_302])])).
fof(f10460, plain, (~ (e20 = op2(e22, e21)) | (~ spl72_101 | spl72_302)), inference(forward_demodulation, [], [f2184, f1190])).
fof(f1190, plain, ((e20 = op2(e21, e22)) | ~ spl72_101), inference(avatar_component_clause, [], [f1188])).
fof(f2184, plain, (~ (op2(e21, e22) = op2(e22, e21)) | spl72_302), inference(avatar_component_clause, [], [f2183])).
fof(f10457, plain, (~ spl72_87 | ~ spl72_107 | spl72_321), inference(avatar_split_clause, [], [f10456, f2276, f1213, f1128])).
fof(f1128, plain, (spl72_87 <=> (e22 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl72_87])])).
fof(f1213, plain, (spl72_107 <=> (e22 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl72_107])])).
fof(f2276, plain, (spl72_321 <=> (e22 = op2(e22, op2(e21, e21)))), introduced(avatar_definition, [new_symbols(naming, [spl72_321])])).
fof(f10456, plain, (~ (e22 = op2(e22, e22)) | (~ spl72_107 | spl72_321)), inference(forward_demodulation, [], [f2278, f1215])).
fof(f1215, plain, ((e22 = op2(e21, e21)) | ~ spl72_107), inference(avatar_component_clause, [], [f1213])).
fof(f2278, plain, (~ (e22 = op2(e22, op2(e21, e21))) | spl72_321), inference(avatar_component_clause, [], [f2276])).
fof(f10455, plain, (~ spl72_87 | ~ spl72_127 | spl72_347), inference(avatar_split_clause, [], [f10454, f2404, f1298, f1128])).
fof(f1298, plain, (spl72_127 <=> (op2(e20, e20) = e22)), introduced(avatar_definition, [new_symbols(naming, [spl72_127])])).
fof(f2404, plain, (spl72_347 <=> (e22 = op2(e22, op2(e20, e20)))), introduced(avatar_definition, [new_symbols(naming, [spl72_347])])).
fof(f10454, plain, (~ (e22 = op2(e22, e22)) | (~ spl72_127 | spl72_347)), inference(forward_demodulation, [], [f2406, f1300])).
fof(f1300, plain, ((op2(e20, e20) = e22) | ~ spl72_127), inference(avatar_component_clause, [], [f1298])).
fof(f2406, plain, (~ (e22 = op2(e22, op2(e20, e20))) | spl72_347), inference(avatar_component_clause, [], [f2404])).
fof(f10449, plain, (~ spl72_100 | spl72_379), inference(avatar_contradiction_clause, [], [f10448])).
fof(f10448, plain, ($false | (~ spl72_100 | spl72_379)), inference(subsumption_resolution, [], [f10447, f1185])).
fof(f1185, plain, ((e23 = op2(e21, e23)) | ~ spl72_100), inference(avatar_component_clause, [], [f1183])).
fof(f10447, plain, (~ (e23 = op2(e21, e23)) | (~ spl72_100 | spl72_379)), inference(forward_demodulation, [], [f2560, f1185])).
fof(f2560, plain, (~ (e23 = op2(e21, op2(e21, e23))) | spl72_379), inference(avatar_component_clause, [], [f2558])).
fof(f2558, plain, (spl72_379 <=> (e23 = op2(e21, op2(e21, e23)))), introduced(avatar_definition, [new_symbols(naming, [spl72_379])])).
fof(f10436, plain, (~ spl72_82 | ~ spl72_415), inference(avatar_split_clause, [], [f10420, f2784, f1107])).
fof(f2784, plain, (spl72_415 <=> (e21 = h4(e12))), introduced(avatar_definition, [new_symbols(naming, [spl72_415])])).
fof(f10420, plain, (~ (e21 = op2(e22, e23)) | ~ spl72_415), inference(backward_demodulation, [], [f2747, f2785])).
fof(f2785, plain, ((e21 = h4(e12)) | ~ spl72_415), inference(avatar_component_clause, [], [f2784])).
fof(f2747, plain, ~ (op2(e22, e23) = h4(e12)), inference(backward_demodulation, [], [f335, f664])).
fof(f335, plain, ~ (op2(e22, e23) = op2(e23, e23)), inference(cnf_transformation, [], [f6])).
fof(f10384, plain, (~ spl72_96 | ~ spl72_490), inference(avatar_split_clause, [], [f10376, f3197, f1166])).
fof(f1166, plain, (spl72_96 <=> (e23 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl72_96])])).
fof(f3197, plain, (spl72_490 <=> (e23 = h3(e12))), introduced(avatar_definition, [new_symbols(naming, [spl72_490])])).
fof(f10376, plain, (~ (e23 = op2(e22, e20)) | ~ spl72_490), inference(backward_demodulation, [], [f2739, f3198])).
fof(f3198, plain, ((e23 = h3(e12)) | ~ spl72_490), inference(avatar_component_clause, [], [f3197])).
fof(f2739, plain, ~ (op2(e22, e20) = h3(e12)), inference(backward_demodulation, [], [f349, f660])).
fof(f660, plain, (op2(e22, e22) = h3(e12)), inference(cnf_transformation, [], [f16])).
fof(f16, plain, ((h3(e13) = op2(e22, op2(e22, op2(e22, e22)))) & (op2(e22, e22) = h3(e12)) & (op2(e22, op2(e22, e22)) = h3(e11)) & (e22 = h3(e10))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG104+1.p', ax16)).
fof(f349, plain, ~ (op2(e22, e20) = op2(e22, e22)), inference(cnf_transformation, [], [f6])).
fof(f10383, plain, (~ spl72_72 | ~ spl72_490), inference(avatar_split_clause, [], [f10375, f3197, f1064])).
fof(f1064, plain, (spl72_72 <=> (e23 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl72_72])])).
fof(f10375, plain, (~ (e23 = op2(e23, e22)) | ~ spl72_490), inference(backward_demodulation, [], [f2738, f3198])).
fof(f2738, plain, ~ (op2(e23, e22) = h3(e12)), inference(backward_demodulation, [], [f329, f660])).
fof(f329, plain, ~ (op2(e22, e22) = op2(e23, e22)), inference(cnf_transformation, [], [f6])).
fof(f10307, plain, (~ spl72_84 | ~ spl72_100), inference(avatar_split_clause, [], [f10306, f1183, f1115])).
fof(f1115, plain, (spl72_84 <=> (e23 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl72_84])])).
fof(f10306, plain, (~ (e23 = op2(e22, e23)) | ~ spl72_100), inference(forward_demodulation, [], [f333, f1185])).
fof(f333, plain, ~ (op2(e21, e23) = op2(e22, e23)), inference(cnf_transformation, [], [f6])).
fof(f10304, plain, (~ spl72_75 | ~ spl72_435), inference(avatar_split_clause, [], [f10303, f2885, f1077])).
fof(f1077, plain, (spl72_75 <=> (e22 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl72_75])])).
fof(f2885, plain, (spl72_435 <=> (e22 = h2(e12))), introduced(avatar_definition, [new_symbols(naming, [spl72_435])])).
fof(f10303, plain, (~ (e22 = op2(e23, e21)) | ~ spl72_435), inference(forward_demodulation, [], [f2729, f2886])).
fof(f2886, plain, ((e22 = h2(e12)) | ~ spl72_435), inference(avatar_component_clause, [], [f2885])).
fof(f2729, plain, ~ (op2(e23, e21) = h2(e12)), inference(backward_demodulation, [], [f322, f656])).
fof(f322, plain, ~ (op2(e21, e21) = op2(e23, e21)), inference(cnf_transformation, [], [f6])).
fof(f10302, plain, (~ spl72_75 | ~ spl72_11 | ~ spl72_447 | ~ spl72_515 | spl72_516), inference(avatar_split_clause, [], [f9920, f3366, f3362, f2946, f773, f1077])).
fof(f773, plain, (spl72_11 <=> (e12 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl72_11])])).
fof(f3366, plain, (spl72_516 <=> (h1(op1(e13, e11)) = op2(h1(e13), e21))), introduced(avatar_definition, [new_symbols(naming, [spl72_516])])).
fof(f9920, plain, (~ (e22 = op2(e23, e21)) | (~ spl72_11 | ~ spl72_447 | ~ spl72_515 | spl72_516)), inference(forward_demodulation, [], [f9919, f2947])).
fof(f9919, plain, (~ (op2(e23, e21) = h1(e12)) | (~ spl72_11 | ~ spl72_515 | spl72_516)), inference(forward_demodulation, [], [f9918, f775])).
fof(f775, plain, ((e12 = op1(e13, e11)) | ~ spl72_11), inference(avatar_component_clause, [], [f773])).
fof(f9918, plain, (~ (op2(e23, e21) = h1(op1(e13, e11))) | (~ spl72_515 | spl72_516)), inference(forward_demodulation, [], [f3368, f3363])).
fof(f3368, plain, (~ (h1(op1(e13, e11)) = op2(h1(e13), e21)) | spl72_516), inference(avatar_component_clause, [], [f3366])).
fof(f10294, plain, (~ spl72_77 | ~ spl72_110 | spl72_284), inference(avatar_contradiction_clause, [], [f10293])).
fof(f10293, plain, ($false | (~ spl72_77 | ~ spl72_110 | spl72_284)), inference(subsumption_resolution, [], [f10292, f1228])).
fof(f1228, plain, ((e21 = op2(e21, e20)) | ~ spl72_110), inference(avatar_component_clause, [], [f1226])).
fof(f1226, plain, (spl72_110 <=> (e21 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl72_110])])).
fof(f10292, plain, (~ (e21 = op2(e21, e20)) | (~ spl72_77 | spl72_284)), inference(forward_demodulation, [], [f2097, f1088])).
fof(f1088, plain, ((e20 = op2(e23, e20)) | ~ spl72_77), inference(avatar_component_clause, [], [f1086])).
fof(f2097, plain, (~ (e21 = op2(e21, op2(e23, e20))) | spl72_284), inference(avatar_component_clause, [], [f2095])).
fof(f2095, plain, (spl72_284 <=> (e21 = op2(e21, op2(e23, e20)))), introduced(avatar_definition, [new_symbols(naming, [spl72_284])])).
fof(f10278, plain, (~ spl72_110 | ~ spl72_113 | spl72_333), inference(avatar_contradiction_clause, [], [f10277])).
fof(f10277, plain, ($false | (~ spl72_110 | ~ spl72_113 | spl72_333)), inference(subsumption_resolution, [], [f10276, f1228])).
fof(f10276, plain, (~ (e21 = op2(e21, e20)) | (~ spl72_113 | spl72_333)), inference(forward_demodulation, [], [f2336, f1241])).
fof(f1241, plain, ((e20 = op2(e20, e23)) | ~ spl72_113), inference(avatar_component_clause, [], [f1239])).
fof(f1239, plain, (spl72_113 <=> (e20 = op2(e20, e23))), introduced(avatar_definition, [new_symbols(naming, [spl72_113])])).
fof(f2336, plain, (~ (e21 = op2(e21, op2(e20, e23))) | spl72_333), inference(avatar_component_clause, [], [f2334])).
fof(f2334, plain, (spl72_333 <=> (e21 = op2(e21, op2(e20, e23)))), introduced(avatar_definition, [new_symbols(naming, [spl72_333])])).
fof(f10235, plain, (~ spl72_74 | spl72_98 | ~ spl72_280), inference(avatar_contradiction_clause, [], [f10234])).
fof(f10234, plain, ($false | (~ spl72_74 | spl72_98 | ~ spl72_280)), inference(subsumption_resolution, [], [f10233, f1176])).
fof(f1176, plain, (~ (e21 = op2(e21, e23)) | spl72_98), inference(avatar_component_clause, [], [f1175])).
fof(f1175, plain, (spl72_98 <=> (e21 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl72_98])])).
fof(f10233, plain, ((e21 = op2(e21, e23)) | (~ spl72_74 | ~ spl72_280)), inference(forward_demodulation, [], [f2078, f1075])).
fof(f1075, plain, ((e21 = op2(e23, e21)) | ~ spl72_74), inference(avatar_component_clause, [], [f1073])).
fof(f2078, plain, ((op2(e21, e23) = op2(e23, e21)) | ~ spl72_280), inference(avatar_component_clause, [], [f2076])).
fof(f2076, plain, (spl72_280 <=> (op2(e21, e23) = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl72_280])])).
fof(f10122, plain, (~ spl72_101 | ~ spl72_110 | spl72_300 | ~ spl72_302), inference(avatar_contradiction_clause, [], [f10121])).
fof(f10121, plain, ($false | (~ spl72_101 | ~ spl72_110 | spl72_300 | ~ spl72_302)), inference(subsumption_resolution, [], [f10120, f1228])).
fof(f10120, plain, (~ (e21 = op2(e21, e20)) | (~ spl72_101 | spl72_300 | ~ spl72_302)), inference(forward_demodulation, [], [f10052, f1190])).
fof(f10052, plain, (~ (e21 = op2(e21, op2(e21, e22))) | (spl72_300 | ~ spl72_302)), inference(backward_demodulation, [], [f2175, f2185])).
fof(f2185, plain, ((op2(e21, e22) = op2(e22, e21)) | ~ spl72_302), inference(avatar_component_clause, [], [f2183])).
fof(f2175, plain, (~ (e21 = op2(e21, op2(e22, e21))) | spl72_300), inference(avatar_component_clause, [], [f2173])).
fof(f2173, plain, (spl72_300 <=> (e21 = op2(e21, op2(e22, e21)))), introduced(avatar_definition, [new_symbols(naming, [spl72_300])])).
fof(f10096, plain, (~ spl72_67 | ~ spl72_96 | ~ spl72_304), inference(avatar_contradiction_clause, [], [f10095])).
fof(f10095, plain, ($false | (~ spl72_67 | ~ spl72_96 | ~ spl72_304)), inference(subsumption_resolution, [], [f10094, f371])).
fof(f371, plain, ~ (e22 = e23), inference(cnf_transformation, [], [f8])).
fof(f8, plain, (~ (e22 = e23) & ~ (e21 = e23) & ~ (e21 = e22) & ~ (e20 = e23) & ~ (e20 = e22) & ~ (e20 = e21)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG104+1.p', ax8)).
fof(f10094, plain, ((e22 = e23) | (~ spl72_67 | ~ spl72_96 | ~ spl72_304)), inference(forward_demodulation, [], [f10087, f1045])).
fof(f1045, plain, ((e22 = op2(e23, e23)) | ~ spl72_67), inference(avatar_component_clause, [], [f1043])).
fof(f1043, plain, (spl72_67 <=> (e22 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl72_67])])).
fof(f10087, plain, ((e23 = op2(e23, e23)) | (~ spl72_96 | ~ spl72_304)), inference(backward_demodulation, [], [f2193, f1168])).
fof(f1168, plain, ((e23 = op2(e22, e20)) | ~ spl72_96), inference(avatar_component_clause, [], [f1166])).
fof(f2193, plain, ((e23 = op2(e23, op2(e22, e20))) | ~ spl72_304), inference(avatar_component_clause, [], [f2192])).
fof(f2192, plain, (spl72_304 <=> (e23 = op2(e23, op2(e22, e20)))), introduced(avatar_definition, [new_symbols(naming, [spl72_304])])).
fof(f10093, plain, (~ spl72_80 | ~ spl72_96), inference(avatar_split_clause, [], [f10086, f1166, f1098])).
fof(f10086, plain, (~ (e23 = op2(e23, e20)) | ~ spl72_96), inference(backward_demodulation, [], [f317, f1168])).
fof(f10081, plain, (~ spl72_101 | ~ spl72_110 | spl72_317), inference(avatar_contradiction_clause, [], [f10080])).
fof(f10080, plain, ($false | (~ spl72_101 | ~ spl72_110 | spl72_317)), inference(subsumption_resolution, [], [f10071, f1228])).
fof(f10071, plain, (~ (e21 = op2(e21, e20)) | (~ spl72_101 | spl72_317)), inference(backward_demodulation, [], [f2258, f1190])).
fof(f2258, plain, (~ (e21 = op2(e21, op2(e21, e22))) | spl72_317), inference(avatar_component_clause, [], [f2256])).
fof(f2256, plain, (spl72_317 <=> (e21 = op2(e21, op2(e21, e22)))), introduced(avatar_definition, [new_symbols(naming, [spl72_317])])).
fof(f10060, plain, (spl72_77 | ~ spl72_113 | ~ spl72_286), inference(avatar_split_clause, [], [f10059, f2105, f1239, f1086])).
fof(f2105, plain, (spl72_286 <=> (op2(e20, e23) = op2(e23, e20))), introduced(avatar_definition, [new_symbols(naming, [spl72_286])])).
fof(f10059, plain, ((e20 = op2(e23, e20)) | (~ spl72_113 | ~ spl72_286)), inference(forward_demodulation, [], [f2107, f1241])).
fof(f2107, plain, ((op2(e20, e23) = op2(e23, e20)) | ~ spl72_286), inference(avatar_component_clause, [], [f2105])).
fof(f10025, plain, (spl72_117 | ~ spl72_296 | ~ spl72_423), inference(avatar_contradiction_clause, [], [f10024])).
fof(f10024, plain, ($false | (spl72_117 | ~ spl72_296 | ~ spl72_423)), inference(subsumption_resolution, [], [f10012, f1257])).
fof(f1257, plain, (~ (e20 = op2(e20, e22)) | spl72_117), inference(avatar_component_clause, [], [f1256])).
fof(f1256, plain, (spl72_117 <=> (e20 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl72_117])])).
fof(f10012, plain, ((e20 = op2(e20, e22)) | (~ spl72_296 | ~ spl72_423)), inference(backward_demodulation, [], [f7945, f2825])).
fof(f2825, plain, ((e22 = h3(e12)) | ~ spl72_423), inference(avatar_component_clause, [], [f2824])).
fof(f2824, plain, (spl72_423 <=> (e22 = h3(e12))), introduced(avatar_definition, [new_symbols(naming, [spl72_423])])).
fof(f7945, plain, ((e20 = op2(e20, h3(e12))) | ~ spl72_296), inference(backward_demodulation, [], [f2155, f660])).
fof(f2155, plain, ((e20 = op2(e20, op2(e22, e22))) | ~ spl72_296), inference(avatar_component_clause, [], [f2154])).
fof(f2154, plain, (spl72_296 <=> (e20 = op2(e20, op2(e22, e22)))), introduced(avatar_definition, [new_symbols(naming, [spl72_296])])).
fof(f10022, plain, (~ spl72_103 | ~ spl72_423), inference(avatar_split_clause, [], [f10007, f2824, f1196])).
fof(f1196, plain, (spl72_103 <=> (e22 = op2(e21, e22))), introduced(avatar_definition, [new_symbols(naming, [spl72_103])])).
fof(f10007, plain, (~ (e22 = op2(e21, e22)) | ~ spl72_423), inference(backward_demodulation, [], [f2737, f2825])).
fof(f2737, plain, ~ (op2(e21, e22) = h3(e12)), inference(backward_demodulation, [], [f327, f660])).
fof(f327, plain, ~ (op2(e21, e22) = op2(e22, e22)), inference(cnf_transformation, [], [f6])).
fof(f9997, plain, (~ spl72_103 | ~ spl72_435), inference(avatar_split_clause, [], [f9987, f2885, f1196])).
fof(f9987, plain, (~ (e22 = op2(e21, e22)) | ~ spl72_435), inference(backward_demodulation, [], [f2731, f2886])).
fof(f2731, plain, ~ (op2(e21, e22) = h2(e12)), inference(backward_demodulation, [], [f345, f656])).
fof(f345, plain, ~ (op2(e21, e21) = op2(e21, e22)), inference(cnf_transformation, [], [f6])).
fof(f9996, plain, (~ spl72_91 | ~ spl72_435), inference(avatar_split_clause, [], [f9986, f2885, f1145])).
fof(f1145, plain, (spl72_91 <=> (e22 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl72_91])])).
fof(f9986, plain, (~ (e22 = op2(e22, e21)) | ~ spl72_435), inference(backward_demodulation, [], [f2728, f2886])).
fof(f9957, plain, (spl72_72 | ~ spl72_411 | ~ spl72_472), inference(avatar_split_clause, [], [f9937, f3091, f2764, f1064])).
fof(f2764, plain, (spl72_411 <=> (e22 = h4(e12))), introduced(avatar_definition, [new_symbols(naming, [spl72_411])])).
fof(f3091, plain, (spl72_472 <=> (e23 = h4(e11))), introduced(avatar_definition, [new_symbols(naming, [spl72_472])])).
fof(f9937, plain, ((e23 = op2(e23, e22)) | (~ spl72_411 | ~ spl72_472)), inference(backward_demodulation, [], [f9825, f3092])).
fof(f3092, plain, ((e23 = h4(e11)) | ~ spl72_472), inference(avatar_component_clause, [], [f3091])).
fof(f9825, plain, ((op2(e23, e22) = h4(e11)) | ~ spl72_411), inference(forward_demodulation, [], [f2752, f2765])).
fof(f2765, plain, ((e22 = h4(e12)) | ~ spl72_411), inference(avatar_component_clause, [], [f2764])).
fof(f2752, plain, (h4(e11) = op2(e23, h4(e12))), inference(forward_demodulation, [], [f663, f664])).
fof(f663, plain, (op2(e23, op2(e23, e23)) = h4(e11)), inference(cnf_transformation, [], [f17])).
fof(f9911, plain, (~ spl72_103 | ~ spl72_39 | ~ spl72_447 | spl72_521), inference(avatar_split_clause, [], [f9910, f3386, f2946, f892, f1196])).
fof(f892, plain, (spl72_39 <=> (e12 = op1(e11, e12))), introduced(avatar_definition, [new_symbols(naming, [spl72_39])])).
fof(f3386, plain, (spl72_521 <=> (h1(op1(e11, e12)) = op2(e21, h1(e12)))), introduced(avatar_definition, [new_symbols(naming, [spl72_521])])).
fof(f9910, plain, (~ (e22 = op2(e21, e22)) | (~ spl72_39 | ~ spl72_447 | spl72_521)), inference(forward_demodulation, [], [f9909, f2947])).
fof(f9909, plain, (~ (op2(e21, e22) = h1(e12)) | (~ spl72_39 | ~ spl72_447 | spl72_521)), inference(forward_demodulation, [], [f9908, f894])).
fof(f894, plain, ((e12 = op1(e11, e12)) | ~ spl72_39), inference(avatar_component_clause, [], [f892])).
fof(f9908, plain, (~ (op2(e21, e22) = h1(op1(e11, e12))) | (~ spl72_447 | spl72_521)), inference(forward_demodulation, [], [f3388, f2947])).
fof(f3388, plain, (~ (h1(op1(e11, e12)) = op2(e21, h1(e12))) | spl72_521), inference(avatar_component_clause, [], [f3386])).
fof(f9814, plain, (spl72_117 | ~ spl72_407 | ~ spl72_411), inference(avatar_contradiction_clause, [], [f9813])).
fof(f9813, plain, ($false | (spl72_117 | ~ spl72_407 | ~ spl72_411)), inference(subsumption_resolution, [], [f9812, f1257])).
fof(f9812, plain, ((e20 = op2(e20, e22)) | (~ spl72_407 | ~ spl72_411)), inference(forward_demodulation, [], [f9556, f2765])).
fof(f9556, plain, ((e20 = op2(e20, h4(e12))) | ~ spl72_407), inference(backward_demodulation, [], [f2696, f664])).
fof(f2696, plain, ((e20 = op2(e20, op2(e23, e23))) | ~ spl72_407), inference(avatar_component_clause, [], [f2695])).
fof(f2695, plain, (spl72_407 <=> (e20 = op2(e20, op2(e23, e23)))), introduced(avatar_definition, [new_symbols(naming, [spl72_407])])).
fof(f9766, plain, (~ spl72_104 | ~ spl72_100), inference(avatar_split_clause, [], [f9765, f1183, f1200])).
fof(f1200, plain, (spl72_104 <=> (e23 = op2(e21, e22))), introduced(avatar_definition, [new_symbols(naming, [spl72_104])])).
fof(f9765, plain, (~ (e23 = op2(e21, e22)) | ~ spl72_100), inference(forward_demodulation, [], [f347, f1185])).
fof(f347, plain, ~ (op2(e21, e22) = op2(e21, e23)), inference(cnf_transformation, [], [f6])).
fof(f9717, plain, (~ spl72_75 | ~ spl72_91), inference(avatar_split_clause, [], [f9708, f1145, f1077])).
fof(f9708, plain, (~ (e22 = op2(e23, e21)) | ~ spl72_91), inference(backward_demodulation, [], [f323, f1147])).
fof(f1147, plain, ((e22 = op2(e22, e21)) | ~ spl72_91), inference(avatar_component_clause, [], [f1145])).
fof(f9673, plain, (~ spl72_103 | spl72_383), inference(avatar_contradiction_clause, [], [f9672])).
fof(f9672, plain, ($false | (~ spl72_103 | spl72_383)), inference(subsumption_resolution, [], [f9671, f1198])).
fof(f1198, plain, ((e22 = op2(e21, e22)) | ~ spl72_103), inference(avatar_component_clause, [], [f1196])).
fof(f9671, plain, (~ (e22 = op2(e21, e22)) | (~ spl72_103 | spl72_383)), inference(forward_demodulation, [], [f2579, f1198])).
fof(f2579, plain, (~ (e22 = op2(e21, op2(e21, e22))) | spl72_383), inference(avatar_component_clause, [], [f2577])).
fof(f2577, plain, (spl72_383 <=> (e22 = op2(e21, op2(e21, e22)))), introduced(avatar_definition, [new_symbols(naming, [spl72_383])])).
fof(f9604, plain, (~ spl72_68 | ~ spl72_4 | spl72_514 | ~ spl72_515), inference(avatar_split_clause, [], [f9603, f3362, f3358, f743, f1047])).
fof(f1047, plain, (spl72_68 <=> (e23 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl72_68])])).
fof(f743, plain, (spl72_4 <=> (e13 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl72_4])])).
fof(f3358, plain, (spl72_514 <=> (h1(op1(e13, e13)) = op2(h1(e13), h1(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl72_514])])).
fof(f9603, plain, (~ (e23 = op2(e23, e23)) | (~ spl72_4 | spl72_514 | ~ spl72_515)), inference(forward_demodulation, [], [f9602, f3363])).
fof(f9602, plain, (~ (op2(e23, e23) = h1(e13)) | (~ spl72_4 | spl72_514 | ~ spl72_515)), inference(forward_demodulation, [], [f9601, f745])).
fof(f745, plain, ((e13 = op1(e13, e13)) | ~ spl72_4), inference(avatar_component_clause, [], [f743])).
fof(f9601, plain, (~ (op2(e23, e23) = h1(op1(e13, e13))) | (spl72_514 | ~ spl72_515)), inference(forward_demodulation, [], [f3360, f3363])).
fof(f3360, plain, (~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | spl72_514), inference(avatar_component_clause, [], [f3358])).
fof(f9517, plain, (~ spl72_99 | ~ spl72_103), inference(avatar_split_clause, [], [f9516, f1196, f1179])).
fof(f9516, plain, (~ (e22 = op2(e21, e23)) | ~ spl72_103), inference(forward_demodulation, [], [f347, f1198])).
fof(f9502, plain, (~ spl72_84 | ~ spl72_490), inference(avatar_split_clause, [], [f9501, f3197, f1115])).
fof(f9501, plain, (~ (e23 = op2(e22, e23)) | ~ spl72_490), inference(forward_demodulation, [], [f2741, f3198])).
fof(f2741, plain, ~ (op2(e22, e23) = h3(e12)), inference(backward_demodulation, [], [f353, f660])).
fof(f353, plain, ~ (op2(e22, e22) = op2(e22, e23)), inference(cnf_transformation, [], [f6])).
fof(f9478, plain, (~ spl72_14 | ~ spl72_78 | ~ spl72_515 | spl72_517), inference(avatar_contradiction_clause, [], [f9477])).
fof(f9477, plain, ($false | (~ spl72_14 | ~ spl72_78 | ~ spl72_515 | spl72_517)), inference(subsumption_resolution, [], [f9476, f1092])).
fof(f1092, plain, ((e21 = op2(e23, e20)) | ~ spl72_78), inference(avatar_component_clause, [], [f1090])).
fof(f9476, plain, (~ (e21 = op2(e23, e20)) | (~ spl72_14 | ~ spl72_515 | spl72_517)), inference(forward_demodulation, [], [f9475, f2727])).
fof(f9475, plain, (~ (op2(e23, e20) = h1(e11)) | (~ spl72_14 | ~ spl72_515 | spl72_517)), inference(forward_demodulation, [], [f9474, f788])).
fof(f788, plain, ((e11 = op1(e13, e10)) | ~ spl72_14), inference(avatar_component_clause, [], [f786])).
fof(f786, plain, (spl72_14 <=> (e11 = op1(e13, e10))), introduced(avatar_definition, [new_symbols(naming, [spl72_14])])).
fof(f9474, plain, (~ (op2(e23, e20) = h1(op1(e13, e10))) | (~ spl72_515 | spl72_517)), inference(forward_demodulation, [], [f3372, f3363])).
fof(f3372, plain, (~ (h1(op1(e13, e10)) = op2(h1(e13), e20)) | spl72_517), inference(avatar_component_clause, [], [f3370])).
fof(f3370, plain, (spl72_517 <=> (h1(op1(e13, e10)) = op2(h1(e13), e20))), introduced(avatar_definition, [new_symbols(naming, [spl72_517])])).
fof(f9445, plain, (~ spl72_29 | ~ spl72_93 | ~ spl72_447 | spl72_519), inference(avatar_contradiction_clause, [], [f9444])).
fof(f9444, plain, ($false | (~ spl72_29 | ~ spl72_93 | ~ spl72_447 | spl72_519)), inference(subsumption_resolution, [], [f9443, f1156])).
fof(f9443, plain, (~ (e20 = op2(e22, e20)) | (~ spl72_29 | ~ spl72_447 | spl72_519)), inference(forward_demodulation, [], [f9442, f650])).
fof(f650, plain, (e20 = h1(e10)), inference(cnf_transformation, [], [f14])).
fof(f9442, plain, (~ (op2(e22, e20) = h1(e10)) | (~ spl72_29 | ~ spl72_447 | spl72_519)), inference(forward_demodulation, [], [f9441, f852])).
fof(f852, plain, ((e10 = op1(e12, e10)) | ~ spl72_29), inference(avatar_component_clause, [], [f850])).
fof(f850, plain, (spl72_29 <=> (e10 = op1(e12, e10))), introduced(avatar_definition, [new_symbols(naming, [spl72_29])])).
fof(f9441, plain, (~ (op2(e22, e20) = h1(op1(e12, e10))) | (~ spl72_447 | spl72_519)), inference(forward_demodulation, [], [f3380, f2947])).
fof(f3380, plain, (~ (h1(op1(e12, e10)) = op2(h1(e12), e20)) | spl72_519), inference(avatar_component_clause, [], [f3378])).
fof(f3378, plain, (spl72_519 <=> (h1(op1(e12, e10)) = op2(h1(e12), e20))), introduced(avatar_definition, [new_symbols(naming, [spl72_519])])).
fof(f9440, plain, (~ spl72_5 | ~ spl72_69 | ~ spl72_447 | spl72_513 | ~ spl72_515), inference(avatar_contradiction_clause, [], [f9439])).
fof(f9439, plain, ($false | (~ spl72_5 | ~ spl72_69 | ~ spl72_447 | spl72_513 | ~ spl72_515)), inference(subsumption_resolution, [], [f9438, f1054])).
fof(f1054, plain, ((e20 = op2(e23, e22)) | ~ spl72_69), inference(avatar_component_clause, [], [f1052])).
fof(f1052, plain, (spl72_69 <=> (e20 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl72_69])])).
fof(f9438, plain, (~ (e20 = op2(e23, e22)) | (~ spl72_5 | ~ spl72_447 | spl72_513 | ~ spl72_515)), inference(forward_demodulation, [], [f9437, f650])).
fof(f9437, plain, (~ (op2(e23, e22) = h1(e10)) | (~ spl72_5 | ~ spl72_447 | spl72_513 | ~ spl72_515)), inference(forward_demodulation, [], [f9436, f750])).
fof(f750, plain, ((e10 = op1(e13, e12)) | ~ spl72_5), inference(avatar_component_clause, [], [f748])).
fof(f748, plain, (spl72_5 <=> (e10 = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl72_5])])).
fof(f9436, plain, (~ (op2(e23, e22) = h1(op1(e13, e12))) | (~ spl72_447 | spl72_513 | ~ spl72_515)), inference(forward_demodulation, [], [f9435, f3363])).
fof(f9435, plain, (~ (h1(op1(e13, e12)) = op2(h1(e13), e22)) | (~ spl72_447 | spl72_513)), inference(forward_demodulation, [], [f3356, f2947])).
fof(f3356, plain, (~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | spl72_513), inference(avatar_component_clause, [], [f3354])).
fof(f3354, plain, (spl72_513 <=> (h1(op1(e13, e12)) = op2(h1(e13), h1(e12)))), introduced(avatar_definition, [new_symbols(naming, [spl72_513])])).
fof(f9427, plain, (~ spl72_34 | ~ spl72_98 | ~ spl72_515 | spl72_520), inference(avatar_contradiction_clause, [], [f9426])).
fof(f9426, plain, ($false | (~ spl72_34 | ~ spl72_98 | ~ spl72_515 | spl72_520)), inference(subsumption_resolution, [], [f9425, f1177])).
fof(f1177, plain, ((e21 = op2(e21, e23)) | ~ spl72_98), inference(avatar_component_clause, [], [f1175])).
fof(f9425, plain, (~ (e21 = op2(e21, e23)) | (~ spl72_34 | ~ spl72_515 | spl72_520)), inference(forward_demodulation, [], [f9424, f2727])).
fof(f9424, plain, (~ (op2(e21, e23) = h1(e11)) | (~ spl72_34 | ~ spl72_515 | spl72_520)), inference(forward_demodulation, [], [f9423, f873])).
fof(f873, plain, ((e11 = op1(e11, e13)) | ~ spl72_34), inference(avatar_component_clause, [], [f871])).
fof(f871, plain, (spl72_34 <=> (e11 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl72_34])])).
fof(f9423, plain, (~ (op2(e21, e23) = h1(op1(e11, e13))) | (~ spl72_515 | spl72_520)), inference(forward_demodulation, [], [f3384, f3363])).
fof(f3384, plain, (~ (h1(op1(e11, e13)) = op2(e21, h1(e13))) | spl72_520), inference(avatar_component_clause, [], [f3382])).
fof(f3382, plain, (spl72_520 <=> (h1(op1(e11, e13)) = op2(e21, h1(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl72_520])])).
fof(f9414, plain, (~ spl72_41 | ~ spl72_443 | spl72_522), inference(avatar_contradiction_clause, [], [f9413])).
fof(f9413, plain, ($false | (~ spl72_41 | ~ spl72_443 | spl72_522)), inference(subsumption_resolution, [], [f9412, f650])).
fof(f9412, plain, (~ (e20 = h1(e10)) | (~ spl72_41 | ~ spl72_443 | spl72_522)), inference(forward_demodulation, [], [f9411, f2927])).
fof(f2927, plain, ((e20 = h2(e12)) | ~ spl72_443), inference(avatar_component_clause, [], [f2926])).
fof(f2926, plain, (spl72_443 <=> (e20 = h2(e12))), introduced(avatar_definition, [new_symbols(naming, [spl72_443])])).
fof(f9411, plain, (~ (h1(e10) = h2(e12)) | (~ spl72_41 | spl72_522)), inference(forward_demodulation, [], [f3392, f903])).
fof(f903, plain, ((e10 = op1(e11, e11)) | ~ spl72_41), inference(avatar_component_clause, [], [f901])).
fof(f901, plain, (spl72_41 <=> (e10 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl72_41])])).
fof(f3392, plain, (~ (h2(e12) = h1(op1(e11, e11))) | spl72_522), inference(avatar_component_clause, [], [f3390])).
fof(f3390, plain, (spl72_522 <=> (h2(e12) = h1(op1(e11, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl72_522])])).
fof(f9405, plain, (~ spl72_63 | spl72_527), inference(avatar_contradiction_clause, [], [f9404])).
fof(f9404, plain, ($false | (~ spl72_63 | spl72_527)), inference(trivial_inequality_removal, [], [f9403])).
fof(f9403, plain, (~ (h1(e12) = h1(e12)) | (~ spl72_63 | spl72_527)), inference(forward_demodulation, [], [f3412, f996])).
fof(f996, plain, ((op1(e10, e10) = e12) | ~ spl72_63), inference(avatar_component_clause, [], [f994])).
fof(f994, plain, (spl72_63 <=> (op1(e10, e10) = e12)), introduced(avatar_definition, [new_symbols(naming, [spl72_63])])).
fof(f3412, plain, (~ (h1(e12) = h1(op1(e10, e10))) | spl72_527), inference(avatar_component_clause, [], [f3410])).
fof(f3410, plain, (spl72_527 <=> (h1(e12) = h1(op1(e10, e10)))), introduced(avatar_definition, [new_symbols(naming, [spl72_527])])).
fof(f9396, plain, (~ spl72_48 | ~ spl72_112 | ~ spl72_515 | spl72_523), inference(avatar_contradiction_clause, [], [f9395])).
fof(f9395, plain, ($false | (~ spl72_48 | ~ spl72_112 | ~ spl72_515 | spl72_523)), inference(subsumption_resolution, [], [f9394, f1236])).
fof(f9394, plain, (~ (e23 = op2(e21, e20)) | (~ spl72_48 | ~ spl72_515 | spl72_523)), inference(forward_demodulation, [], [f9393, f3363])).
fof(f9393, plain, (~ (op2(e21, e20) = h1(e13)) | (~ spl72_48 | spl72_523)), inference(forward_demodulation, [], [f3396, f932])).
fof(f932, plain, ((e13 = op1(e11, e10)) | ~ spl72_48), inference(avatar_component_clause, [], [f930])).
fof(f930, plain, (spl72_48 <=> (e13 = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl72_48])])).
fof(f3396, plain, (~ (op2(e21, e20) = h1(op1(e11, e10))) | spl72_523), inference(avatar_component_clause, [], [f3394])).
fof(f3394, plain, (spl72_523 <=> (op2(e21, e20) = h1(op1(e11, e10)))), introduced(avatar_definition, [new_symbols(naming, [spl72_523])])).
fof(f9354, plain, (~ spl72_54 | spl72_525), inference(avatar_contradiction_clause, [], [f9353])).
fof(f9353, plain, ($false | (~ spl72_54 | spl72_525)), inference(subsumption_resolution, [], [f9352, f2727])).
fof(f9352, plain, (~ (e21 = h1(e11)) | (~ spl72_54 | spl72_525)), inference(forward_demodulation, [], [f3404, f958])).
fof(f958, plain, ((e11 = op1(e10, e12)) | ~ spl72_54), inference(avatar_component_clause, [], [f956])).
fof(f956, plain, (spl72_54 <=> (e11 = op1(e10, e12))), introduced(avatar_definition, [new_symbols(naming, [spl72_54])])).
fof(f3404, plain, (~ (e21 = h1(op1(e10, e12))) | spl72_525), inference(avatar_component_clause, [], [f3402])).
fof(f3402, plain, (spl72_525 <=> (e21 = h1(op1(e10, e12)))), introduced(avatar_definition, [new_symbols(naming, [spl72_525])])).
fof(f9317, plain, (~ spl72_18 | ~ spl72_26), inference(avatar_split_clause, [], [f9306, f837, f803])).
fof(f803, plain, (spl72_18 <=> (e11 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl72_18])])).
fof(f9306, plain, (~ (e11 = op1(e12, e13)) | ~ spl72_26), inference(backward_demodulation, [], [f304, f839])).
fof(f304, plain, ~ (op1(e12, e11) = op1(e12, e13)), inference(cnf_transformation, [], [f5])).
fof(f5, plain, (~ (op1(e13, e12) = op1(e13, e13)) & ~ (op1(e13, e11) = op1(e13, e13)) & ~ (op1(e13, e11) = op1(e13, e12)) & ~ (op1(e13, e10) = op1(e13, e13)) & ~ (op1(e13, e10) = op1(e13, e12)) & ~ (op1(e13, e10) = op1(e13, e11)) & ~ (op1(e12, e12) = op1(e12, e13)) & ~ (op1(e12, e11) = op1(e12, e13)) & ~ (op1(e12, e11) = op1(e12, e12)) & ~ (op1(e12, e10) = op1(e12, e13)) & ~ (op1(e12, e10) = op1(e12, e12)) & ~ (op1(e12, e10) = op1(e12, e11)) & ~ (op1(e11, e12) = op1(e11, e13)) & ~ (op1(e11, e11) = op1(e11, e13)) & ~ (op1(e11, e11) = op1(e11, e12)) & ~ (op1(e11, e10) = op1(e11, e13)) & ~ (op1(e11, e10) = op1(e11, e12)) & ~ (op1(e11, e10) = op1(e11, e11)) & ~ (op1(e10, e12) = op1(e10, e13)) & ~ (op1(e10, e11) = op1(e10, e13)) & ~ (op1(e10, e11) = op1(e10, e12)) & ~ (op1(e10, e10) = op1(e10, e13)) & ~ (op1(e10, e10) = op1(e10, e12)) & ~ (op1(e10, e10) = op1(e10, e11)) & ~ (op1(e12, e13) = op1(e13, e13)) & ~ (op1(e11, e13) = op1(e13, e13)) & ~ (op1(e11, e13) = op1(e12, e13)) & ~ (op1(e10, e13) = op1(e13, e13)) & ~ (op1(e10, e13) = op1(e12, e13)) & ~ (op1(e10, e13) = op1(e11, e13)) & ~ (op1(e12, e12) = op1(e13, e12)) & ~ (op1(e11, e12) = op1(e13, e12)) & ~ (op1(e11, e12) = op1(e12, e12)) & ~ (op1(e10, e12) = op1(e13, e12)) & ~ (op1(e10, e12) = op1(e12, e12)) & ~ (op1(e10, e12) = op1(e11, e12)) & ~ (op1(e12, e11) = op1(e13, e11)) & ~ (op1(e11, e11) = op1(e13, e11)) & ~ (op1(e11, e11) = op1(e12, e11)) & ~ (op1(e10, e11) = op1(e13, e11)) & ~ (op1(e10, e11) = op1(e12, e11)) & ~ (op1(e10, e11) = op1(e11, e11)) & ~ (op1(e12, e10) = op1(e13, e10)) & ~ (op1(e11, e10) = op1(e13, e10)) & ~ (op1(e11, e10) = op1(e12, e10)) & ~ (op1(e10, e10) = op1(e13, e10)) & ~ (op1(e10, e10) = op1(e12, e10)) & ~ (op1(e10, e10) = op1(e11, e10))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG104+1.p', ax5)).
fof(f9304, plain, (~ spl72_25 | ~ spl72_29), inference(avatar_split_clause, [], [f9293, f850, f833])).
fof(f833, plain, (spl72_25 <=> (e10 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl72_25])])).
fof(f9293, plain, (~ (e10 = op1(e12, e11)) | ~ spl72_29), inference(backward_demodulation, [], [f300, f852])).
fof(f300, plain, ~ (op1(e12, e10) = op1(e12, e11)), inference(cnf_transformation, [], [f5])).
fof(f9288, plain, (~ spl72_18 | ~ spl72_34), inference(avatar_split_clause, [], [f9275, f871, f803])).
fof(f9275, plain, (~ (e11 = op1(e12, e13)) | ~ spl72_34), inference(backward_demodulation, [], [f285, f873])).
fof(f285, plain, ~ (op1(e11, e13) = op1(e12, e13)), inference(cnf_transformation, [], [f5])).
fof(f9267, plain, (~ spl72_35 | ~ spl72_39), inference(avatar_split_clause, [], [f9252, f892, f875])).
fof(f875, plain, (spl72_35 <=> (e12 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl72_35])])).
fof(f9252, plain, (~ (e12 = op1(e11, e13)) | ~ spl72_39), inference(backward_demodulation, [], [f299, f894])).
fof(f299, plain, ~ (op1(e11, e12) = op1(e11, e13)), inference(cnf_transformation, [], [f5])).
fof(f9266, plain, (~ spl72_23 | ~ spl72_39), inference(avatar_split_clause, [], [f9251, f892, f824])).
fof(f824, plain, (spl72_23 <=> (e12 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl72_23])])).
fof(f9251, plain, (~ (e12 = op1(e12, e12)) | ~ spl72_39), inference(backward_demodulation, [], [f279, f894])).
fof(f279, plain, ~ (op1(e11, e12) = op1(e12, e12)), inference(cnf_transformation, [], [f5])).
fof(f9245, plain, (~ spl72_25 | ~ spl72_41), inference(avatar_split_clause, [], [f9233, f901, f833])).
fof(f9233, plain, (~ (e10 = op1(e12, e11)) | ~ spl72_41), inference(backward_demodulation, [], [f273, f903])).
fof(f273, plain, ~ (op1(e11, e11) = op1(e12, e11)), inference(cnf_transformation, [], [f5])).
fof(f9231, plain, (~ spl72_40 | ~ spl72_48), inference(avatar_split_clause, [], [f9221, f930, f896])).
fof(f896, plain, (spl72_40 <=> (e13 = op1(e11, e12))), introduced(avatar_definition, [new_symbols(naming, [spl72_40])])).
fof(f9221, plain, (~ (e13 = op1(e11, e12)) | ~ spl72_48), inference(backward_demodulation, [], [f295, f932])).
fof(f295, plain, ~ (op1(e11, e10) = op1(e11, e12)), inference(cnf_transformation, [], [f5])).
fof(f9230, plain, (~ spl72_32 | ~ spl72_48), inference(avatar_split_clause, [], [f9219, f930, f862])).
fof(f862, plain, (spl72_32 <=> (e13 = op1(e12, e10))), introduced(avatar_definition, [new_symbols(naming, [spl72_32])])).
fof(f9219, plain, (~ (e13 = op1(e12, e10)) | ~ spl72_48), inference(backward_demodulation, [], [f267, f932])).
fof(f267, plain, ~ (op1(e11, e10) = op1(e12, e10)), inference(cnf_transformation, [], [f5])).
fof(f9214, plain, (spl72_48 | ~ spl72_60 | ~ spl72_189), inference(avatar_split_clause, [], [f9213, f1631, f981, f930])).
fof(f981, plain, (spl72_60 <=> (e13 = op1(e10, e11))), introduced(avatar_definition, [new_symbols(naming, [spl72_60])])).
fof(f1631, plain, (spl72_189 <=> (op1(e10, e11) = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl72_189])])).
fof(f9213, plain, ((e13 = op1(e11, e10)) | (~ spl72_60 | ~ spl72_189)), inference(forward_demodulation, [], [f1633, f983])).
fof(f983, plain, ((e13 = op1(e10, e11)) | ~ spl72_60), inference(avatar_component_clause, [], [f981])).
fof(f1633, plain, ((op1(e10, e11) = op1(e11, e10)) | ~ spl72_189), inference(avatar_component_clause, [], [f1631])).
fof(f9208, plain, (~ spl72_23 | ~ spl72_63 | spl72_207), inference(avatar_split_clause, [], [f9207, f1720, f994, f824])).
fof(f1720, plain, (spl72_207 <=> (e12 = op1(e12, op1(e10, e10)))), introduced(avatar_definition, [new_symbols(naming, [spl72_207])])).
fof(f9207, plain, (~ (e12 = op1(e12, e12)) | (~ spl72_63 | spl72_207)), inference(forward_demodulation, [], [f1722, f996])).
fof(f1722, plain, (~ (e12 = op1(e12, op1(e10, e10))) | spl72_207), inference(avatar_component_clause, [], [f1720])).
fof(f9130, plain, (~ spl72_7 | ~ spl72_11), inference(avatar_split_clause, [], [f9118, f773, f756])).
fof(f756, plain, (spl72_7 <=> (e12 = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl72_7])])).
fof(f9118, plain, (~ (e12 = op1(e13, e12)) | ~ spl72_11), inference(backward_demodulation, [], [f309, f775])).
fof(f309, plain, ~ (op1(e13, e11) = op1(e13, e12)), inference(cnf_transformation, [], [f5])).
fof(f9109, plain, (~ spl72_25 | ~ spl72_113 | ~ spl72_443 | ~ spl72_495 | ~ spl72_510), inference(avatar_contradiction_clause, [], [f9108])).
fof(f9108, plain, ($false | (~ spl72_25 | ~ spl72_113 | ~ spl72_443 | ~ spl72_495 | ~ spl72_510)), inference(subsumption_resolution, [], [f9107, f366])).
fof(f366, plain, ~ (e20 = e21), inference(cnf_transformation, [], [f8])).
fof(f9107, plain, ((e20 = e21) | (~ spl72_25 | ~ spl72_113 | ~ spl72_443 | ~ spl72_495 | ~ spl72_510)), inference(backward_demodulation, [], [f654, f9101])).
fof(f9101, plain, ((e20 = h2(e10)) | (~ spl72_25 | ~ spl72_113 | ~ spl72_443 | ~ spl72_495 | ~ spl72_510)), inference(backward_demodulation, [], [f8689, f835])).
fof(f835, plain, ((e10 = op1(e12, e11)) | ~ spl72_25), inference(avatar_component_clause, [], [f833])).
fof(f8689, plain, ((e20 = h2(op1(e12, e11))) | (~ spl72_113 | ~ spl72_443 | ~ spl72_495 | ~ spl72_510)), inference(forward_demodulation, [], [f8688, f1241])).
fof(f8688, plain, ((op2(e20, e23) = h2(op1(e12, e11))) | (~ spl72_443 | ~ spl72_495 | ~ spl72_510)), inference(forward_demodulation, [], [f8687, f2927])).
fof(f8687, plain, ((h2(op1(e12, e11)) = op2(h2(e12), e23)) | (~ spl72_495 | ~ spl72_510)), inference(forward_demodulation, [], [f3241, f3323])).
fof(f3323, plain, ((e23 = h2(e11)) | ~ spl72_510), inference(avatar_component_clause, [], [f3322])).
fof(f3322, plain, (spl72_510 <=> (e23 = h2(e11))), introduced(avatar_definition, [new_symbols(naming, [spl72_510])])).
fof(f3241, plain, ((h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ spl72_495), inference(avatar_component_clause, [], [f3240])).
fof(f3240, plain, (spl72_495 <=> (h2(op1(e12, e11)) = op2(h2(e12), h2(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl72_495])])).
fof(f654, plain, (e21 = h2(e10)), inference(cnf_transformation, [], [f15])).
fof(f9105, plain, (~ spl72_21 | ~ spl72_25), inference(avatar_split_clause, [], [f9091, f833, f816])).
fof(f816, plain, (spl72_21 <=> (e10 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl72_21])])).
fof(f9091, plain, (~ (e10 = op1(e12, e12)) | ~ spl72_25), inference(backward_demodulation, [], [f303, f835])).
fof(f303, plain, ~ (op1(e12, e11) = op1(e12, e12)), inference(cnf_transformation, [], [f5])).
fof(f9104, plain, (~ spl72_9 | ~ spl72_25), inference(avatar_split_clause, [], [f9090, f833, f765])).
fof(f765, plain, (spl72_9 <=> (e10 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl72_9])])).
fof(f9090, plain, (~ (e10 = op1(e13, e11)) | ~ spl72_25), inference(backward_demodulation, [], [f275, f835])).
fof(f275, plain, ~ (op1(e12, e11) = op1(e13, e11)), inference(cnf_transformation, [], [f5])).
fof(f9079, plain, (~ spl72_42 | spl72_182), inference(avatar_contradiction_clause, [], [f9078])).
fof(f9078, plain, ($false | (~ spl72_42 | spl72_182)), inference(subsumption_resolution, [], [f9077, f907])).
fof(f907, plain, ((e11 = op1(e11, e11)) | ~ spl72_42), inference(avatar_component_clause, [], [f905])).
fof(f905, plain, (spl72_42 <=> (e11 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl72_42])])).
fof(f9077, plain, (~ (e11 = op1(e11, e11)) | (~ spl72_42 | spl72_182)), inference(forward_demodulation, [], [f1599, f907])).
fof(f1599, plain, (~ (e11 = op1(e11, op1(e11, e11))) | spl72_182), inference(avatar_component_clause, [], [f1597])).
fof(f1597, plain, (spl72_182 <=> (e11 = op1(e11, op1(e11, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl72_182])])).
fof(f9067, plain, (~ spl72_9 | ~ spl72_14 | spl72_221), inference(avatar_split_clause, [], [f9066, f1787, f786, f765])).
fof(f1787, plain, (spl72_221 <=> (e10 = op1(e13, op1(e13, e10)))), introduced(avatar_definition, [new_symbols(naming, [spl72_221])])).
fof(f9066, plain, (~ (e10 = op1(e13, e11)) | (~ spl72_14 | spl72_221)), inference(forward_demodulation, [], [f1789, f788])).
fof(f1789, plain, (~ (e10 = op1(e13, op1(e13, e10))) | spl72_221), inference(avatar_component_clause, [], [f1787])).
fof(f9056, plain, (~ spl72_9 | ~ spl72_60 | spl72_259), inference(avatar_split_clause, [], [f9055, f1970, f981, f765])).
fof(f1970, plain, (spl72_259 <=> (e10 = op1(op1(e10, e11), e11))), introduced(avatar_definition, [new_symbols(naming, [spl72_259])])).
fof(f9055, plain, (~ (e10 = op1(e13, e11)) | (~ spl72_60 | spl72_259)), inference(forward_demodulation, [], [f1972, f983])).
fof(f1972, plain, (~ (e10 = op1(op1(e10, e11), e11)) | spl72_259), inference(avatar_component_clause, [], [f1970])).
fof(f9029, plain, (~ spl72_8 | ~ spl72_40), inference(avatar_split_clause, [], [f9028, f896, f760])).
fof(f760, plain, (spl72_8 <=> (e13 = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl72_8])])).
fof(f9028, plain, (~ (e13 = op1(e13, e12)) | ~ spl72_40), inference(forward_demodulation, [], [f280, f898])).
fof(f898, plain, ((e13 = op1(e11, e12)) | ~ spl72_40), inference(avatar_component_clause, [], [f896])).
fof(f280, plain, ~ (op1(e11, e12) = op1(e13, e12)), inference(cnf_transformation, [], [f5])).
fof(f9009, plain, (~ spl72_24 | ~ spl72_32), inference(avatar_split_clause, [], [f9003, f862, f828])).
fof(f828, plain, (spl72_24 <=> (e13 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl72_24])])).
fof(f9003, plain, (~ (e13 = op1(e12, e12)) | ~ spl72_32), inference(backward_demodulation, [], [f301, f864])).
fof(f864, plain, ((e13 = op1(e12, e10)) | ~ spl72_32), inference(avatar_component_clause, [], [f862])).
fof(f301, plain, ~ (op1(e12, e10) = op1(e12, e12)), inference(cnf_transformation, [], [f5])).
fof(f9000, plain, (~ spl72_35 | ~ spl72_40 | spl72_243), inference(avatar_contradiction_clause, [], [f8999])).
fof(f8999, plain, ($false | (~ spl72_35 | ~ spl72_40 | spl72_243)), inference(subsumption_resolution, [], [f8988, f877])).
fof(f877, plain, ((e12 = op1(e11, e13)) | ~ spl72_35), inference(avatar_component_clause, [], [f875])).
fof(f8988, plain, (~ (e12 = op1(e11, e13)) | (~ spl72_40 | spl72_243)), inference(backward_demodulation, [], [f1895, f898])).
fof(f1895, plain, (~ (e12 = op1(e11, op1(e11, e12))) | spl72_243), inference(avatar_component_clause, [], [f1893])).
fof(f1893, plain, (spl72_243 <=> (e12 = op1(e11, op1(e11, e12)))), introduced(avatar_definition, [new_symbols(naming, [spl72_243])])).
fof(f8997, plain, (spl72_6 | ~ spl72_40 | ~ spl72_241), inference(avatar_contradiction_clause, [], [f8996])).
fof(f8996, plain, ($false | (spl72_6 | ~ spl72_40 | ~ spl72_241)), inference(subsumption_resolution, [], [f8986, f753])).
fof(f753, plain, (~ (e11 = op1(e13, e12)) | spl72_6), inference(avatar_component_clause, [], [f752])).
fof(f752, plain, (spl72_6 <=> (e11 = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl72_6])])).
fof(f8986, plain, ((e11 = op1(e13, e12)) | (~ spl72_40 | ~ spl72_241)), inference(backward_demodulation, [], [f1884, f898])).
fof(f1884, plain, ((e11 = op1(op1(e11, e12), e12)) | ~ spl72_241), inference(avatar_component_clause, [], [f1883])).
fof(f1883, plain, (spl72_241 <=> (e11 = op1(op1(e11, e12), e12))), introduced(avatar_definition, [new_symbols(naming, [spl72_241])])).
fof(f8992, plain, (~ spl72_24 | ~ spl72_40), inference(avatar_split_clause, [], [f8981, f896, f828])).
fof(f8981, plain, (~ (e13 = op1(e12, e12)) | ~ spl72_40), inference(backward_demodulation, [], [f279, f898])).
fof(f8977, plain, (~ spl72_45 | spl72_249), inference(avatar_contradiction_clause, [], [f8976])).
fof(f8976, plain, ($false | (~ spl72_45 | spl72_249)), inference(subsumption_resolution, [], [f8963, f920])).
fof(f920, plain, ((e10 = op1(e11, e10)) | ~ spl72_45), inference(avatar_component_clause, [], [f918])).
fof(f918, plain, (spl72_45 <=> (e10 = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl72_45])])).
fof(f8963, plain, (~ (e10 = op1(e11, e10)) | (~ spl72_45 | spl72_249)), inference(backward_demodulation, [], [f1925, f920])).
fof(f1925, plain, (~ (e10 = op1(e11, op1(e11, e10))) | spl72_249), inference(avatar_component_clause, [], [f1923])).
fof(f1923, plain, (spl72_249 <=> (e10 = op1(e11, op1(e11, e10)))), introduced(avatar_definition, [new_symbols(naming, [spl72_249])])).
fof(f8966, plain, (~ spl72_37 | ~ spl72_45), inference(avatar_split_clause, [], [f8954, f918, f884])).
fof(f884, plain, (spl72_37 <=> (e10 = op1(e11, e12))), introduced(avatar_definition, [new_symbols(naming, [spl72_37])])).
fof(f8954, plain, (~ (e10 = op1(e11, e12)) | ~ spl72_45), inference(backward_demodulation, [], [f295, f920])).
fof(f8952, plain, (~ spl72_48 | ~ spl72_60 | spl72_189), inference(avatar_split_clause, [], [f8951, f1631, f981, f930])).
fof(f8951, plain, (~ (e13 = op1(e11, e10)) | (~ spl72_60 | spl72_189)), inference(forward_demodulation, [], [f1632, f983])).
fof(f1632, plain, (~ (op1(e10, e11) = op1(e11, e10)) | spl72_189), inference(avatar_component_clause, [], [f1631])).
fof(f8948, plain, (spl72_40 | ~ spl72_18 | ~ spl72_224), inference(avatar_split_clause, [], [f8947, f1801, f803, f896])).
fof(f1801, plain, (spl72_224 <=> (e13 = op1(op1(e12, e13), e12))), introduced(avatar_definition, [new_symbols(naming, [spl72_224])])).
fof(f8947, plain, ((e13 = op1(e11, e12)) | (~ spl72_18 | ~ spl72_224)), inference(forward_demodulation, [], [f1803, f805])).
fof(f805, plain, ((e11 = op1(e12, e13)) | ~ spl72_18), inference(avatar_component_clause, [], [f803])).
fof(f1803, plain, ((e13 = op1(op1(e12, e13), e12)) | ~ spl72_224), inference(avatar_component_clause, [], [f1801])).
fof(f8927, plain, (~ spl72_40 | ~ spl72_35 | spl72_239), inference(avatar_split_clause, [], [f8878, f1874, f875, f896])).
fof(f1874, plain, (spl72_239 <=> (e13 = op1(e11, op1(e11, e13)))), introduced(avatar_definition, [new_symbols(naming, [spl72_239])])).
fof(f8878, plain, (~ (e13 = op1(e11, e12)) | (~ spl72_35 | spl72_239)), inference(forward_demodulation, [], [f1876, f877])).
fof(f1876, plain, (~ (e13 = op1(e11, op1(e11, e13))) | spl72_239), inference(avatar_component_clause, [], [f1874])).
fof(f8917, plain, (~ spl72_14 | ~ spl72_48 | spl72_247), inference(avatar_contradiction_clause, [], [f8916])).
fof(f8916, plain, ($false | (~ spl72_14 | ~ spl72_48 | spl72_247)), inference(subsumption_resolution, [], [f8915, f788])).
fof(f8915, plain, (~ (e11 = op1(e13, e10)) | (~ spl72_48 | spl72_247)), inference(forward_demodulation, [], [f1915, f932])).
fof(f1915, plain, (~ (e11 = op1(op1(e11, e10), e10)) | spl72_247), inference(avatar_component_clause, [], [f1913])).
fof(f1913, plain, (spl72_247 <=> (e11 = op1(op1(e11, e10), e10))), introduced(avatar_definition, [new_symbols(naming, [spl72_247])])).
fof(f8886, plain, (~ spl72_18 | ~ spl72_35 | spl72_237), inference(avatar_contradiction_clause, [], [f8885])).
fof(f8885, plain, ($false | (~ spl72_18 | ~ spl72_35 | spl72_237)), inference(subsumption_resolution, [], [f8884, f805])).
fof(f8884, plain, (~ (e11 = op1(e12, e13)) | (~ spl72_35 | spl72_237)), inference(forward_demodulation, [], [f1866, f877])).
fof(f1866, plain, (~ (e11 = op1(op1(e11, e13), e13)) | spl72_237), inference(avatar_component_clause, [], [f1864])).
fof(f1864, plain, (spl72_237 <=> (e11 = op1(op1(e11, e13), e13))), introduced(avatar_definition, [new_symbols(naming, [spl72_237])])).
fof(f8883, plain, (~ spl72_27 | ~ spl72_35 | ~ spl72_238), inference(avatar_contradiction_clause, [], [f8882])).
fof(f8882, plain, ($false | (~ spl72_27 | ~ spl72_35 | ~ spl72_238)), inference(subsumption_resolution, [], [f8881, f365])).
fof(f365, plain, ~ (e12 = e13), inference(cnf_transformation, [], [f7])).
fof(f7, plain, (~ (e12 = e13) & ~ (e11 = e13) & ~ (e11 = e12) & ~ (e10 = e13) & ~ (e10 = e12) & ~ (e10 = e11)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG104+1.p', ax7)).
fof(f8881, plain, ((e12 = e13) | (~ spl72_27 | ~ spl72_35 | ~ spl72_238)), inference(forward_demodulation, [], [f8880, f843])).
fof(f843, plain, ((e12 = op1(e12, e11)) | ~ spl72_27), inference(avatar_component_clause, [], [f841])).
fof(f841, plain, (spl72_27 <=> (e12 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl72_27])])).
fof(f8880, plain, ((e13 = op1(e12, e11)) | (~ spl72_35 | ~ spl72_238)), inference(forward_demodulation, [], [f1871, f877])).
fof(f1871, plain, ((e13 = op1(op1(e11, e13), e11)) | ~ spl72_238), inference(avatar_component_clause, [], [f1869])).
fof(f1869, plain, (spl72_238 <=> (e13 = op1(op1(e11, e13), e11))), introduced(avatar_definition, [new_symbols(naming, [spl72_238])])).
fof(f8836, plain, (~ spl72_9 | ~ spl72_49 | ~ spl72_216), inference(avatar_contradiction_clause, [], [f8835])).
fof(f8835, plain, ($false | (~ spl72_9 | ~ spl72_49 | ~ spl72_216)), inference(subsumption_resolution, [], [f8834, f360])).
fof(f360, plain, ~ (e10 = e11), inference(cnf_transformation, [], [f7])).
fof(f8834, plain, ((e10 = e11) | (~ spl72_9 | ~ spl72_49 | ~ spl72_216)), inference(forward_demodulation, [], [f8827, f937])).
fof(f937, plain, ((e10 = op1(e10, e13)) | ~ spl72_49), inference(avatar_component_clause, [], [f935])).
fof(f935, plain, (spl72_49 <=> (e10 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl72_49])])).
fof(f8827, plain, ((e11 = op1(e10, e13)) | (~ spl72_9 | ~ spl72_216)), inference(backward_demodulation, [], [f1765, f767])).
fof(f767, plain, ((e10 = op1(e13, e11)) | ~ spl72_9), inference(avatar_component_clause, [], [f765])).
fof(f1765, plain, ((e11 = op1(op1(e13, e11), e13)) | ~ spl72_216), inference(avatar_component_clause, [], [f1763])).
fof(f1763, plain, (spl72_216 <=> (e11 = op1(op1(e13, e11), e13))), introduced(avatar_definition, [new_symbols(naming, [spl72_216])])).
fof(f8830, plain, (~ spl72_5 | ~ spl72_9), inference(avatar_split_clause, [], [f8819, f765, f748])).
fof(f8819, plain, (~ (e10 = op1(e13, e12)) | ~ spl72_9), inference(backward_demodulation, [], [f309, f767])).
fof(f8816, plain, (~ spl72_24 | ~ spl72_447 | ~ spl72_490 | spl72_511 | ~ spl72_515), inference(avatar_contradiction_clause, [], [f8815])).
fof(f8815, plain, ($false | (~ spl72_24 | ~ spl72_447 | ~ spl72_490 | spl72_511 | ~ spl72_515)), inference(subsumption_resolution, [], [f8810, f3363])).
fof(f8810, plain, (~ (e23 = h1(e13)) | (~ spl72_24 | ~ spl72_447 | ~ spl72_490 | spl72_511)), inference(backward_demodulation, [], [f8449, f830])).
fof(f830, plain, ((e13 = op1(e12, e12)) | ~ spl72_24), inference(avatar_component_clause, [], [f828])).
fof(f8449, plain, (~ (e23 = h1(op1(e12, e12))) | (~ spl72_447 | ~ spl72_490 | spl72_511)), inference(forward_demodulation, [], [f7946, f3198])).
fof(f7946, plain, (~ (h3(e12) = h1(op1(e12, e12))) | (~ spl72_447 | spl72_511)), inference(backward_demodulation, [], [f7647, f660])).
fof(f7647, plain, (~ (op2(e22, e22) = h1(op1(e12, e12))) | (~ spl72_447 | spl72_511)), inference(forward_demodulation, [], [f3348, f2947])).
fof(f3348, plain, (~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | spl72_511), inference(avatar_component_clause, [], [f3346])).
fof(f3346, plain, (spl72_511 <=> (h1(op1(e12, e12)) = op2(h1(e12), h1(e12)))), introduced(avatar_definition, [new_symbols(naming, [spl72_511])])).
fof(f8802, plain, (~ spl72_23 | ~ spl72_27), inference(avatar_split_clause, [], [f8792, f841, f824])).
fof(f8792, plain, (~ (e12 = op1(e12, e12)) | ~ spl72_27), inference(backward_demodulation, [], [f303, f843])).
fof(f8801, plain, (~ spl72_11 | ~ spl72_27), inference(avatar_split_clause, [], [f8791, f841, f773])).
fof(f8791, plain, (~ (e12 = op1(e13, e11)) | ~ spl72_27), inference(backward_demodulation, [], [f275, f843])).
fof(f8790, plain, (~ spl72_29 | spl72_235), inference(avatar_contradiction_clause, [], [f8789])).
fof(f8789, plain, ($false | (~ spl72_29 | spl72_235)), inference(subsumption_resolution, [], [f8781, f852])).
fof(f8781, plain, (~ (e10 = op1(e12, e10)) | (~ spl72_29 | spl72_235)), inference(backward_demodulation, [], [f1857, f852])).
fof(f1857, plain, (~ (e10 = op1(e12, op1(e12, e10))) | spl72_235), inference(avatar_component_clause, [], [f1855])).
fof(f1855, plain, (spl72_235 <=> (e10 = op1(e12, op1(e12, e10)))), introduced(avatar_definition, [new_symbols(naming, [spl72_235])])).
fof(f8763, plain, (~ spl72_25 | ~ spl72_37 | spl72_162), inference(avatar_split_clause, [], [f8751, f1499, f884, f833])).
fof(f1499, plain, (spl72_162 <=> (op1(e11, e12) = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl72_162])])).
fof(f8751, plain, (~ (e10 = op1(e12, e11)) | (~ spl72_37 | spl72_162)), inference(backward_demodulation, [], [f1500, f886])).
fof(f886, plain, ((e10 = op1(e11, e12)) | ~ spl72_37), inference(avatar_component_clause, [], [f884])).
fof(f1500, plain, (~ (op1(e11, e12) = op1(e12, e11)) | spl72_162), inference(avatar_component_clause, [], [f1499])).
fof(f8762, plain, (~ spl72_5 | ~ spl72_37), inference(avatar_split_clause, [], [f8747, f884, f748])).
fof(f8747, plain, (~ (e10 = op1(e13, e12)) | ~ spl72_37), inference(backward_demodulation, [], [f280, f886])).
fof(f8725, plain, (~ spl72_23 | ~ spl72_35 | spl72_171), inference(avatar_split_clause, [], [f8724, f1542, f875, f824])).
fof(f1542, plain, (spl72_171 <=> (e12 = op1(e12, op1(e11, e13)))), introduced(avatar_definition, [new_symbols(naming, [spl72_171])])).
fof(f8724, plain, (~ (e12 = op1(e12, e12)) | (~ spl72_35 | spl72_171)), inference(forward_demodulation, [], [f1544, f877])).
fof(f1544, plain, (~ (e12 = op1(e12, op1(e11, e13))) | spl72_171), inference(avatar_component_clause, [], [f1542])).
fof(f8612, plain, (~ spl72_18 | ~ spl72_35 | spl72_223), inference(avatar_contradiction_clause, [], [f8611])).
fof(f8611, plain, ($false | (~ spl72_18 | ~ spl72_35 | spl72_223)), inference(subsumption_resolution, [], [f8604, f877])).
fof(f8604, plain, (~ (e12 = op1(e11, e13)) | (~ spl72_18 | spl72_223)), inference(backward_demodulation, [], [f1798, f805])).
fof(f1798, plain, (~ (e12 = op1(op1(e12, e13), e13)) | spl72_223), inference(avatar_component_clause, [], [f1796])).
fof(f1796, plain, (spl72_223 <=> (e12 = op1(op1(e12, e13), e13))), introduced(avatar_definition, [new_symbols(naming, [spl72_223])])).
fof(f8578, plain, (spl72_5 | ~ spl72_32 | ~ spl72_234), inference(avatar_split_clause, [], [f8573, f1850, f862, f748])).
fof(f1850, plain, (spl72_234 <=> (e10 = op1(op1(e12, e10), e12))), introduced(avatar_definition, [new_symbols(naming, [spl72_234])])).
fof(f8573, plain, ((e10 = op1(e13, e12)) | (~ spl72_32 | ~ spl72_234)), inference(backward_demodulation, [], [f1852, f864])).
fof(f1852, plain, ((e10 = op1(op1(e12, e10), e12)) | ~ spl72_234), inference(avatar_component_clause, [], [f1850])).
fof(f8577, plain, (~ spl72_20 | ~ spl72_32), inference(avatar_split_clause, [], [f8571, f862, f811])).
fof(f811, plain, (spl72_20 <=> (e13 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl72_20])])).
fof(f8571, plain, (~ (e13 = op1(e12, e13)) | ~ spl72_32), inference(backward_demodulation, [], [f302, f864])).
fof(f302, plain, ~ (op1(e12, e10) = op1(e12, e13)), inference(cnf_transformation, [], [f5])).
fof(f8537, plain, (spl72_11 | ~ spl72_40 | ~ spl72_242), inference(avatar_split_clause, [], [f8527, f1888, f896, f773])).
fof(f1888, plain, (spl72_242 <=> (e12 = op1(op1(e11, e12), e11))), introduced(avatar_definition, [new_symbols(naming, [spl72_242])])).
fof(f8527, plain, ((e12 = op1(e13, e11)) | (~ spl72_40 | ~ spl72_242)), inference(backward_demodulation, [], [f1890, f898])).
fof(f1890, plain, ((e12 = op1(op1(e11, e12), e11)) | ~ spl72_242), inference(avatar_component_clause, [], [f1888])).
fof(f8536, plain, (~ spl72_40 | ~ spl72_49 | spl72_178), inference(avatar_contradiction_clause, [], [f8535])).
fof(f8535, plain, ($false | (~ spl72_40 | ~ spl72_49 | spl72_178)), inference(subsumption_resolution, [], [f8525, f937])).
fof(f8525, plain, (~ (e10 = op1(e10, e13)) | (~ spl72_40 | spl72_178)), inference(backward_demodulation, [], [f1579, f898])).
fof(f1579, plain, (~ (e10 = op1(e10, op1(e11, e12))) | spl72_178), inference(avatar_component_clause, [], [f1577])).
fof(f1577, plain, (spl72_178 <=> (e10 = op1(e10, op1(e11, e12)))), introduced(avatar_definition, [new_symbols(naming, [spl72_178])])).
fof(f8516, plain, (~ spl72_34 | ~ spl72_42), inference(avatar_split_clause, [], [f8508, f905, f871])).
fof(f8508, plain, (~ (e11 = op1(e11, e13)) | ~ spl72_42), inference(backward_demodulation, [], [f298, f907])).
fof(f298, plain, ~ (op1(e11, e11) = op1(e11, e13)), inference(cnf_transformation, [], [f5])).
fof(f8515, plain, (~ spl72_10 | ~ spl72_42), inference(avatar_split_clause, [], [f8506, f905, f769])).
fof(f769, plain, (spl72_10 <=> (e11 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl72_10])])).
fof(f8506, plain, (~ (e11 = op1(e13, e11)) | ~ spl72_42), inference(backward_demodulation, [], [f274, f907])).
fof(f274, plain, ~ (op1(e11, e11) = op1(e13, e11)), inference(cnf_transformation, [], [f5])).
fof(f8503, plain, (~ spl72_41 | ~ spl72_45), inference(avatar_split_clause, [], [f8493, f918, f901])).
fof(f8493, plain, (~ (e10 = op1(e11, e11)) | ~ spl72_45), inference(backward_demodulation, [], [f294, f920])).
fof(f294, plain, ~ (op1(e11, e10) = op1(e11, e11)), inference(cnf_transformation, [], [f5])).
fof(f8502, plain, (~ spl72_13 | ~ spl72_45), inference(avatar_split_clause, [], [f8492, f918, f782])).
fof(f782, plain, (spl72_13 <=> (e10 = op1(e13, e10))), introduced(avatar_definition, [new_symbols(naming, [spl72_13])])).
fof(f8492, plain, (~ (e10 = op1(e13, e10)) | ~ spl72_45), inference(backward_demodulation, [], [f268, f920])).
fof(f268, plain, ~ (op1(e11, e10) = op1(e13, e10)), inference(cnf_transformation, [], [f5])).
fof(f8488, plain, (spl72_14 | ~ spl72_60 | ~ spl72_260), inference(avatar_split_clause, [], [f8487, f1975, f981, f786])).
fof(f1975, plain, (spl72_260 <=> (e11 = op1(op1(e10, e11), e10))), introduced(avatar_definition, [new_symbols(naming, [spl72_260])])).
fof(f8487, plain, ((e11 = op1(e13, e10)) | (~ spl72_60 | ~ spl72_260)), inference(forward_demodulation, [], [f1977, f983])).
fof(f1977, plain, ((e11 = op1(op1(e10, e11), e10)) | ~ spl72_260), inference(avatar_component_clause, [], [f1975])).
fof(f8456, plain, (~ spl72_42 | ~ spl72_54 | spl72_198), inference(avatar_split_clause, [], [f6882, f1675, f956, f905])).
fof(f1675, plain, (spl72_198 <=> (e11 = op1(e11, op1(e10, e12)))), introduced(avatar_definition, [new_symbols(naming, [spl72_198])])).
fof(f6882, plain, (~ (e11 = op1(e11, e11)) | (~ spl72_54 | spl72_198)), inference(forward_demodulation, [], [f1677, f958])).
fof(f1677, plain, (~ (e11 = op1(e11, op1(e10, e12))) | spl72_198), inference(avatar_component_clause, [], [f1675])).
fof(f8444, plain, (spl72_13 | ~ spl72_49 | ~ spl72_146), inference(avatar_split_clause, [], [f8443, f1421, f935, f782])).
fof(f1421, plain, (spl72_146 <=> (op1(e10, e13) = op1(e13, e10))), introduced(avatar_definition, [new_symbols(naming, [spl72_146])])).
fof(f8443, plain, ((e10 = op1(e13, e10)) | (~ spl72_49 | ~ spl72_146)), inference(forward_demodulation, [], [f1423, f937])).
fof(f1423, plain, ((op1(e10, e13) = op1(e13, e10)) | ~ spl72_146), inference(avatar_component_clause, [], [f1421])).
fof(f8405, plain, (~ spl72_18 | ~ spl72_30), inference(avatar_split_clause, [], [f8397, f854, f803])).
fof(f854, plain, (spl72_30 <=> (e11 = op1(e12, e10))), introduced(avatar_definition, [new_symbols(naming, [spl72_30])])).
fof(f8397, plain, (~ (e11 = op1(e12, e13)) | ~ spl72_30), inference(backward_demodulation, [], [f302, f856])).
fof(f856, plain, ((e11 = op1(e12, e10)) | ~ spl72_30), inference(avatar_component_clause, [], [f854])).
fof(f8360, plain, (~ spl72_36 | ~ spl72_48), inference(avatar_split_clause, [], [f8352, f930, f879])).
fof(f879, plain, (spl72_36 <=> (e13 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl72_36])])).
fof(f8352, plain, (~ (e13 = op1(e11, e13)) | ~ spl72_48), inference(backward_demodulation, [], [f296, f932])).
fof(f296, plain, ~ (op1(e11, e10) = op1(e11, e13)), inference(cnf_transformation, [], [f5])).
fof(f8331, plain, (~ spl72_71 | ~ spl72_75), inference(avatar_split_clause, [], [f8322, f1077, f1060])).
fof(f1060, plain, (spl72_71 <=> (e22 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl72_71])])).
fof(f8322, plain, (~ (e22 = op2(e23, e22)) | ~ spl72_75), inference(backward_demodulation, [], [f357, f1079])).
fof(f1079, plain, ((e22 = op2(e23, e21)) | ~ spl72_75), inference(avatar_component_clause, [], [f1077])).
fof(f357, plain, ~ (op2(e23, e21) = op2(e23, e22)), inference(cnf_transformation, [], [f6])).
fof(f8312, plain, (~ spl72_71 | ~ spl72_103), inference(avatar_split_clause, [], [f8309, f1196, f1060])).
fof(f8309, plain, (~ (e22 = op2(e23, e22)) | ~ spl72_103), inference(backward_demodulation, [], [f328, f1198])).
fof(f328, plain, ~ (op2(e21, e22) = op2(e23, e22)), inference(cnf_transformation, [], [f6])).
fof(f8305, plain, (spl72_510 | ~ spl72_112 | ~ spl72_443), inference(avatar_split_clause, [], [f8304, f2926, f1234, f3322])).
fof(f8304, plain, ((e23 = h2(e11)) | (~ spl72_112 | ~ spl72_443)), inference(backward_demodulation, [], [f8224, f1236])).
fof(f8224, plain, ((op2(e21, e20) = h2(e11)) | ~ spl72_443), inference(backward_demodulation, [], [f2734, f2927])).
fof(f2734, plain, (h2(e11) = op2(e21, h2(e12))), inference(forward_demodulation, [], [f655, f656])).
fof(f655, plain, (op2(e21, op2(e21, e21)) = h2(e11)), inference(cnf_transformation, [], [f15])).
fof(f8300, plain, (spl72_34 | ~ spl72_10 | ~ spl72_140), inference(avatar_split_clause, [], [f8299, f1392, f769, f871])).
fof(f1392, plain, (spl72_140 <=> (op1(e11, e13) = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl72_140])])).
fof(f8299, plain, ((e11 = op1(e11, e13)) | (~ spl72_10 | ~ spl72_140)), inference(forward_demodulation, [], [f1394, f771])).
fof(f771, plain, ((e11 = op1(e13, e11)) | ~ spl72_10), inference(avatar_component_clause, [], [f769])).
fof(f1394, plain, ((op1(e11, e13) = op1(e13, e11)) | ~ spl72_140), inference(avatar_component_clause, [], [f1392])).
fof(f8293, plain, (~ spl72_46 | ~ spl72_49 | spl72_193), inference(avatar_split_clause, [], [f8292, f1650, f935, f922])).
fof(f922, plain, (spl72_46 <=> (e11 = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl72_46])])).
fof(f1650, plain, (spl72_193 <=> (e11 = op1(e11, op1(e10, e13)))), introduced(avatar_definition, [new_symbols(naming, [spl72_193])])).
fof(f8292, plain, (~ (e11 = op1(e11, e10)) | (~ spl72_49 | spl72_193)), inference(forward_demodulation, [], [f1652, f937])).
fof(f1652, plain, (~ (e11 = op1(e11, op1(e10, e13))) | spl72_193), inference(avatar_component_clause, [], [f1650])).
fof(f8269, plain, (~ spl72_73 | ~ spl72_124 | spl72_399), inference(avatar_split_clause, [], [f8268, f2654, f1285, f1069])).
fof(f1069, plain, (spl72_73 <=> (e20 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl72_73])])).
fof(f2654, plain, (spl72_399 <=> (e20 = op2(op2(e20, e21), e21))), introduced(avatar_definition, [new_symbols(naming, [spl72_399])])).
fof(f8268, plain, (~ (e20 = op2(e23, e21)) | (~ spl72_124 | spl72_399)), inference(forward_demodulation, [], [f2656, f1287])).
fof(f2656, plain, (~ (e20 = op2(op2(e20, e21), e21)) | spl72_399), inference(avatar_component_clause, [], [f2654])).
fof(f8252, plain, (~ spl72_27 | ~ spl72_328 | ~ spl72_443 | spl72_495), inference(avatar_contradiction_clause, [], [f8251])).
fof(f8251, plain, ($false | (~ spl72_27 | ~ spl72_328 | ~ spl72_443 | spl72_495)), inference(subsumption_resolution, [], [f8234, f8242])).
fof(f8242, plain, ((e20 = op2(e20, h2(e11))) | (~ spl72_328 | ~ spl72_443)), inference(backward_demodulation, [], [f2311, f8224])).
fof(f2311, plain, ((e20 = op2(e20, op2(e21, e20))) | ~ spl72_328), inference(avatar_component_clause, [], [f2310])).
fof(f2310, plain, (spl72_328 <=> (e20 = op2(e20, op2(e21, e20)))), introduced(avatar_definition, [new_symbols(naming, [spl72_328])])).
fof(f8234, plain, (~ (e20 = op2(e20, h2(e11))) | (~ spl72_27 | ~ spl72_443 | spl72_495)), inference(backward_demodulation, [], [f8177, f2927])).
fof(f8177, plain, (~ (h2(e12) = op2(h2(e12), h2(e11))) | (~ spl72_27 | spl72_495)), inference(forward_demodulation, [], [f3242, f843])).
fof(f3242, plain, (~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | spl72_495), inference(avatar_component_clause, [], [f3240])).
fof(f8237, plain, (~ spl72_109 | ~ spl72_443), inference(avatar_split_clause, [], [f8222, f2926, f1222])).
fof(f1222, plain, (spl72_109 <=> (e20 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl72_109])])).
fof(f8222, plain, (~ (e20 = op2(e21, e20)) | ~ spl72_443), inference(backward_demodulation, [], [f2730, f2927])).
fof(f2730, plain, ~ (op2(e21, e20) = h2(e12)), inference(backward_demodulation, [], [f342, f656])).
fof(f342, plain, ~ (op2(e21, e20) = op2(e21, e21)), inference(cnf_transformation, [], [f6])).
fof(f8236, plain, (~ spl72_73 | ~ spl72_443), inference(avatar_split_clause, [], [f8221, f2926, f1069])).
fof(f8221, plain, (~ (e20 = op2(e23, e21)) | ~ spl72_443), inference(backward_demodulation, [], [f2729, f2927])).
fof(f8200, plain, (~ spl72_104 | ~ spl72_490), inference(avatar_split_clause, [], [f8187, f3197, f1200])).
fof(f8187, plain, (~ (e23 = op2(e21, e22)) | ~ spl72_490), inference(backward_demodulation, [], [f2737, f3198])).
fof(f8150, plain, (~ spl72_27 | spl72_229), inference(avatar_contradiction_clause, [], [f8149])).
fof(f8149, plain, ($false | (~ spl72_27 | spl72_229)), inference(subsumption_resolution, [], [f8148, f843])).
fof(f8148, plain, (~ (e12 = op1(e12, e11)) | (~ spl72_27 | spl72_229)), inference(forward_demodulation, [], [f1828, f843])).
fof(f1828, plain, (~ (e12 = op1(op1(e12, e11), e11)) | spl72_229), inference(avatar_component_clause, [], [f1826])).
fof(f1826, plain, (spl72_229 <=> (e12 = op1(op1(e12, e11), e11))), introduced(avatar_definition, [new_symbols(naming, [spl72_229])])).
fof(f8132, plain, (~ spl72_74 | ~ spl72_78), inference(avatar_split_clause, [], [f8131, f1090, f1073])).
fof(f8131, plain, (~ (e21 = op2(e23, e21)) | ~ spl72_78), inference(forward_demodulation, [], [f354, f1092])).
fof(f354, plain, ~ (op2(e23, e20) = op2(e23, e21)), inference(cnf_transformation, [], [f6])).
fof(f8058, plain, (~ spl72_37 | ~ spl72_41), inference(avatar_split_clause, [], [f8046, f901, f884])).
fof(f8046, plain, (~ (e10 = op1(e11, e12)) | ~ spl72_41), inference(backward_demodulation, [], [f297, f903])).
fof(f297, plain, ~ (op1(e11, e11) = op1(e11, e12)), inference(cnf_transformation, [], [f5])).
fof(f8032, plain, (spl72_70 | ~ spl72_104 | ~ spl72_381), inference(avatar_contradiction_clause, [], [f8031])).
fof(f8031, plain, ($false | (spl72_70 | ~ spl72_104 | ~ spl72_381)), inference(subsumption_resolution, [], [f8027, f1057])).
fof(f1057, plain, (~ (e21 = op2(e23, e22)) | spl72_70), inference(avatar_component_clause, [], [f1056])).
fof(f1056, plain, (spl72_70 <=> (e21 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl72_70])])).
fof(f8027, plain, ((e21 = op2(e23, e22)) | (~ spl72_104 | ~ spl72_381)), inference(backward_demodulation, [], [f2568, f1202])).
fof(f1202, plain, ((e23 = op2(e21, e22)) | ~ spl72_104), inference(avatar_component_clause, [], [f1200])).
fof(f2568, plain, ((e21 = op2(op2(e21, e22), e22)) | ~ spl72_381), inference(avatar_component_clause, [], [f2567])).
fof(f2567, plain, (spl72_381 <=> (e21 = op2(op2(e21, e22), e22))), introduced(avatar_definition, [new_symbols(naming, [spl72_381])])).
fof(f8008, plain, (~ spl72_112 | ~ spl72_78 | spl72_359), inference(avatar_split_clause, [], [f8007, f2461, f1090, f1234])).
fof(f2461, plain, (spl72_359 <=> (e23 = op2(op2(e23, e20), e20))), introduced(avatar_definition, [new_symbols(naming, [spl72_359])])).
fof(f8007, plain, (~ (e23 = op2(e21, e20)) | (~ spl72_78 | spl72_359)), inference(forward_demodulation, [], [f2463, f1092])).
fof(f2463, plain, (~ (e23 = op2(op2(e23, e20), e20)) | spl72_359), inference(avatar_component_clause, [], [f2461])).
fof(f7991, plain, (~ spl72_101 | ~ spl72_431), inference(avatar_split_clause, [], [f7975, f2865, f1188])).
fof(f2865, plain, (spl72_431 <=> (e20 = h3(e12))), introduced(avatar_definition, [new_symbols(naming, [spl72_431])])).
fof(f7975, plain, (~ (e20 = op2(e21, e22)) | ~ spl72_431), inference(backward_demodulation, [], [f2737, f2866])).
fof(f2866, plain, ((e20 = h3(e12)) | ~ spl72_431), inference(avatar_component_clause, [], [f2865])).
fof(f7931, plain, (spl72_88 | ~ spl72_83 | ~ spl72_364), inference(avatar_split_clause, [], [f7930, f2485, f1111, f1132])).
fof(f2485, plain, (spl72_364 <=> (e23 = op2(op2(e22, e23), e22))), introduced(avatar_definition, [new_symbols(naming, [spl72_364])])).
fof(f7930, plain, ((e23 = op2(e22, e22)) | (~ spl72_83 | ~ spl72_364)), inference(forward_demodulation, [], [f2487, f1113])).
fof(f1113, plain, ((e22 = op2(e22, e23)) | ~ spl72_83), inference(avatar_component_clause, [], [f1111])).
fof(f2487, plain, ((e23 = op2(op2(e22, e23), e22)) | ~ spl72_364), inference(avatar_component_clause, [], [f2485])).
fof(f7929, plain, (~ spl72_83 | spl72_363), inference(avatar_contradiction_clause, [], [f7928])).
fof(f7928, plain, ($false | (~ spl72_83 | spl72_363)), inference(subsumption_resolution, [], [f7927, f1113])).
fof(f7927, plain, (~ (e22 = op2(e22, e23)) | (~ spl72_83 | spl72_363)), inference(forward_demodulation, [], [f2482, f1113])).
fof(f2482, plain, (~ (e22 = op2(op2(e22, e23), e23)) | spl72_363), inference(avatar_component_clause, [], [f2480])).
fof(f2480, plain, (spl72_363 <=> (e22 = op2(op2(e22, e23), e23))), introduced(avatar_definition, [new_symbols(naming, [spl72_363])])).
fof(f7893, plain, (~ spl72_94 | ~ spl72_78), inference(avatar_split_clause, [], [f7510, f1090, f1158])).
fof(f1158, plain, (spl72_94 <=> (e21 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl72_94])])).
fof(f7510, plain, (~ (e21 = op2(e22, e20)) | ~ spl72_78), inference(forward_demodulation, [], [f317, f1092])).
fof(f7892, plain, (~ spl72_94 | ~ spl72_90), inference(avatar_split_clause, [], [f7891, f1141, f1158])).
fof(f7891, plain, (~ (e21 = op2(e22, e20)) | ~ spl72_90), inference(forward_demodulation, [], [f348, f1143])).
fof(f7772, plain, (~ spl72_58 | ~ spl72_60), inference(avatar_contradiction_clause, [], [f7771])).
fof(f7771, plain, ($false | (~ spl72_58 | ~ spl72_60)), inference(subsumption_resolution, [], [f7770, f364])).
fof(f364, plain, ~ (e11 = e13), inference(cnf_transformation, [], [f7])).
fof(f7770, plain, ((e11 = e13) | (~ spl72_58 | ~ spl72_60)), inference(backward_demodulation, [], [f983, f975])).
fof(f975, plain, ((e11 = op1(e10, e11)) | ~ spl72_58), inference(avatar_component_clause, [], [f973])).
fof(f973, plain, (spl72_58 <=> (e11 = op1(e10, e11))), introduced(avatar_definition, [new_symbols(naming, [spl72_58])])).
fof(f7730, plain, (spl72_58 | ~ spl72_46 | ~ spl72_189), inference(avatar_split_clause, [], [f7729, f1631, f922, f973])).
fof(f7729, plain, ((e11 = op1(e10, e11)) | (~ spl72_46 | ~ spl72_189)), inference(forward_demodulation, [], [f1633, f924])).
fof(f924, plain, ((e11 = op1(e11, e10)) | ~ spl72_46), inference(avatar_component_clause, [], [f922])).
fof(f7609, plain, (~ spl72_83 | ~ spl72_68 | spl72_405), inference(avatar_split_clause, [], [f6960, f2685, f1047, f1111])).
fof(f2685, plain, (spl72_405 <=> (e22 = op2(e22, op2(e23, e23)))), introduced(avatar_definition, [new_symbols(naming, [spl72_405])])).
fof(f6960, plain, (~ (e22 = op2(e22, e23)) | (~ spl72_68 | spl72_405)), inference(backward_demodulation, [], [f2687, f1049])).
fof(f1049, plain, ((e23 = op2(e23, e23)) | ~ spl72_68), inference(avatar_component_clause, [], [f1047])).
fof(f2687, plain, (~ (e22 = op2(e22, op2(e23, e23))) | spl72_405), inference(avatar_component_clause, [], [f2685])).
fof(f7592, plain, (~ spl72_23 | ~ spl72_43 | spl72_181), inference(avatar_contradiction_clause, [], [f7591])).
fof(f7591, plain, ($false | (~ spl72_23 | ~ spl72_43 | spl72_181)), inference(subsumption_resolution, [], [f7590, f826])).
fof(f826, plain, ((e12 = op1(e12, e12)) | ~ spl72_23), inference(avatar_component_clause, [], [f824])).
fof(f7590, plain, (~ (e12 = op1(e12, e12)) | (~ spl72_43 | spl72_181)), inference(forward_demodulation, [], [f1594, f911])).
fof(f911, plain, ((e12 = op1(e11, e11)) | ~ spl72_43), inference(avatar_component_clause, [], [f909])).
fof(f909, plain, (spl72_43 <=> (e12 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl72_43])])).
fof(f1594, plain, (~ (e12 = op1(e12, op1(e11, e11))) | spl72_181), inference(avatar_component_clause, [], [f1592])).
fof(f1592, plain, (spl72_181 <=> (e12 = op1(e12, op1(e11, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl72_181])])).
fof(f7578, plain, (~ spl72_78 | ~ spl72_112 | spl72_387), inference(avatar_contradiction_clause, [], [f7577])).
fof(f7577, plain, ($false | (~ spl72_78 | ~ spl72_112 | spl72_387)), inference(subsumption_resolution, [], [f7576, f1092])).
fof(f7576, plain, (~ (e21 = op2(e23, e20)) | (~ spl72_112 | spl72_387)), inference(forward_demodulation, [], [f2599, f1236])).
fof(f2599, plain, (~ (e21 = op2(op2(e21, e20), e20)) | spl72_387), inference(avatar_component_clause, [], [f2597])).
fof(f2597, plain, (spl72_387 <=> (e21 = op2(op2(e21, e20), e20))), introduced(avatar_definition, [new_symbols(naming, [spl72_387])])).
fof(f7528, plain, (~ spl72_49 | ~ spl72_113 | ~ spl72_515 | spl72_524), inference(avatar_contradiction_clause, [], [f7527])).
fof(f7527, plain, ($false | (~ spl72_49 | ~ spl72_113 | ~ spl72_515 | spl72_524)), inference(subsumption_resolution, [], [f7526, f1241])).
fof(f7526, plain, (~ (e20 = op2(e20, e23)) | (~ spl72_49 | ~ spl72_515 | spl72_524)), inference(forward_demodulation, [], [f7525, f650])).
fof(f7525, plain, (~ (op2(e20, e23) = h1(e10)) | (~ spl72_49 | ~ spl72_515 | spl72_524)), inference(forward_demodulation, [], [f7524, f937])).
fof(f7524, plain, (~ (op2(e20, e23) = h1(op1(e10, e13))) | (~ spl72_515 | spl72_524)), inference(forward_demodulation, [], [f3400, f3363])).
fof(f3400, plain, (~ (h1(op1(e10, e13)) = op2(e20, h1(e13))) | spl72_524), inference(avatar_component_clause, [], [f3398])).
fof(f3398, plain, (spl72_524 <=> (h1(op1(e10, e13)) = op2(e20, h1(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl72_524])])).
fof(f7515, plain, (~ spl72_113 | spl72_296 | ~ spl72_490), inference(avatar_contradiction_clause, [], [f7514])).
fof(f7514, plain, ($false | (~ spl72_113 | spl72_296 | ~ spl72_490)), inference(subsumption_resolution, [], [f7513, f1241])).
fof(f7513, plain, (~ (e20 = op2(e20, e23)) | (spl72_296 | ~ spl72_490)), inference(forward_demodulation, [], [f5722, f3198])).
fof(f5722, plain, (~ (e20 = op2(e20, h3(e12))) | spl72_296), inference(backward_demodulation, [], [f2156, f660])).
fof(f2156, plain, (~ (e20 = op2(e20, op2(e22, e22))) | spl72_296), inference(avatar_component_clause, [], [f2154])).
fof(f7476, plain, (~ spl72_101 | spl72_111 | ~ spl72_383), inference(avatar_contradiction_clause, [], [f7475])).
fof(f7475, plain, ($false | (~ spl72_101 | spl72_111 | ~ spl72_383)), inference(subsumption_resolution, [], [f7472, f1231])).
fof(f1231, plain, (~ (e22 = op2(e21, e20)) | spl72_111), inference(avatar_component_clause, [], [f1230])).
fof(f1230, plain, (spl72_111 <=> (e22 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl72_111])])).
fof(f7472, plain, ((e22 = op2(e21, e20)) | (~ spl72_101 | ~ spl72_383)), inference(backward_demodulation, [], [f2578, f1190])).
fof(f2578, plain, ((e22 = op2(e21, op2(e21, e22))) | ~ spl72_383), inference(avatar_component_clause, [], [f2577])).
fof(f7474, plain, (~ spl72_101 | ~ spl72_118 | spl72_381), inference(avatar_contradiction_clause, [], [f7473])).
fof(f7473, plain, ($false | (~ spl72_101 | ~ spl72_118 | spl72_381)), inference(subsumption_resolution, [], [f7470, f1262])).
fof(f1262, plain, ((e21 = op2(e20, e22)) | ~ spl72_118), inference(avatar_component_clause, [], [f1260])).
fof(f1260, plain, (spl72_118 <=> (e21 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl72_118])])).
fof(f7470, plain, (~ (e21 = op2(e20, e22)) | (~ spl72_101 | spl72_381)), inference(backward_demodulation, [], [f2569, f1190])).
fof(f2569, plain, (~ (e21 = op2(op2(e21, e22), e22)) | spl72_381), inference(avatar_component_clause, [], [f2567])).
fof(f7459, plain, (~ spl72_104 | ~ spl72_112), inference(avatar_split_clause, [], [f7452, f1234, f1200])).
fof(f7452, plain, (~ (e23 = op2(e21, e22)) | ~ spl72_112), inference(backward_demodulation, [], [f343, f1236])).
fof(f343, plain, ~ (op2(e21, e20) = op2(e21, e22)), inference(cnf_transformation, [], [f6])).
fof(f7458, plain, (~ spl72_96 | ~ spl72_112), inference(avatar_split_clause, [], [f7451, f1234, f1166])).
fof(f7451, plain, (~ (e23 = op2(e22, e20)) | ~ spl72_112), inference(backward_demodulation, [], [f315, f1236])).
fof(f315, plain, ~ (op2(e21, e20) = op2(e22, e20)), inference(cnf_transformation, [], [f6])).
fof(f7447, plain, (~ spl72_13 | ~ spl72_46 | spl72_144), inference(avatar_contradiction_clause, [], [f7446])).
fof(f7446, plain, ($false | (~ spl72_13 | ~ spl72_46 | spl72_144)), inference(subsumption_resolution, [], [f7445, f924])).
fof(f7445, plain, (~ (e11 = op1(e11, e10)) | (~ spl72_13 | spl72_144)), inference(forward_demodulation, [], [f1413, f784])).
fof(f784, plain, ((e10 = op1(e13, e10)) | ~ spl72_13), inference(avatar_component_clause, [], [f782])).
fof(f1413, plain, (~ (e11 = op1(e11, op1(e13, e10))) | spl72_144), inference(avatar_component_clause, [], [f1411])).
fof(f1411, plain, (spl72_144 <=> (e11 = op1(e11, op1(e13, e10)))), introduced(avatar_definition, [new_symbols(naming, [spl72_144])])).
fof(f7429, plain, (spl72_111 | ~ spl72_118 | ~ spl72_396), inference(avatar_contradiction_clause, [], [f7428])).
fof(f7428, plain, ($false | (spl72_111 | ~ spl72_118 | ~ spl72_396)), inference(subsumption_resolution, [], [f7427, f1231])).
fof(f7427, plain, ((e22 = op2(e21, e20)) | (~ spl72_118 | ~ spl72_396)), inference(forward_demodulation, [], [f2642, f1262])).
fof(f2642, plain, ((e22 = op2(op2(e20, e22), e20)) | ~ spl72_396), inference(avatar_component_clause, [], [f2640])).
fof(f2640, plain, (spl72_396 <=> (e22 = op2(op2(e20, e22), e20))), introduced(avatar_definition, [new_symbols(naming, [spl72_396])])).
fof(f7334, plain, (spl72_104 | ~ spl72_82 | ~ spl72_364), inference(avatar_split_clause, [], [f7070, f2485, f1107, f1200])).
fof(f7070, plain, ((e23 = op2(e21, e22)) | (~ spl72_82 | ~ spl72_364)), inference(forward_demodulation, [], [f2487, f1109])).
fof(f1109, plain, ((e21 = op2(e22, e23)) | ~ spl72_82), inference(avatar_component_clause, [], [f1107])).
fof(f7329, plain, (~ spl72_88 | ~ spl72_71 | spl72_351), inference(avatar_split_clause, [], [f7138, f2423, f1060, f1132])).
fof(f2423, plain, (spl72_351 <=> (e23 = op2(op2(e23, e22), e22))), introduced(avatar_definition, [new_symbols(naming, [spl72_351])])).
fof(f7138, plain, (~ (e23 = op2(e22, e22)) | (~ spl72_71 | spl72_351)), inference(backward_demodulation, [], [f2425, f1062])).
fof(f1062, plain, ((e22 = op2(e23, e22)) | ~ spl72_71), inference(avatar_component_clause, [], [f1060])).
fof(f2425, plain, (~ (e23 = op2(op2(e23, e22), e22)) | spl72_351), inference(avatar_component_clause, [], [f2423])).
fof(f7318, plain, (~ spl72_73 | ~ spl72_77), inference(avatar_split_clause, [], [f7309, f1086, f1069])).
fof(f7309, plain, (~ (e20 = op2(e23, e21)) | ~ spl72_77), inference(backward_demodulation, [], [f354, f1088])).
fof(f7300, plain, (~ spl72_78 | ~ spl72_110), inference(avatar_split_clause, [], [f7291, f1226, f1090])).
fof(f7291, plain, (~ (e21 = op2(e23, e20)) | ~ spl72_110), inference(backward_demodulation, [], [f316, f1228])).
fof(f7271, plain, (~ spl72_82 | ~ spl72_99 | spl72_363), inference(avatar_contradiction_clause, [], [f7270])).
fof(f7270, plain, ($false | (~ spl72_82 | ~ spl72_99 | spl72_363)), inference(subsumption_resolution, [], [f7269, f1181])).
fof(f7269, plain, (~ (e22 = op2(e21, e23)) | (~ spl72_82 | spl72_363)), inference(forward_demodulation, [], [f2482, f1109])).
fof(f7172, plain, (spl72_95 | ~ spl72_294 | ~ spl72_431), inference(avatar_contradiction_clause, [], [f7171])).
fof(f7171, plain, ($false | (spl72_95 | ~ spl72_294 | ~ spl72_431)), inference(subsumption_resolution, [], [f7170, f1163])).
fof(f1163, plain, (~ (e22 = op2(e22, e20)) | spl72_95), inference(avatar_component_clause, [], [f1162])).
fof(f1162, plain, (spl72_95 <=> (e22 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl72_95])])).
fof(f7170, plain, ((e22 = op2(e22, e20)) | (~ spl72_294 | ~ spl72_431)), inference(forward_demodulation, [], [f7011, f2866])).
fof(f7011, plain, ((e22 = op2(e22, h3(e12))) | ~ spl72_294), inference(backward_demodulation, [], [f2145, f660])).
fof(f7133, plain, (~ spl72_73 | ~ spl72_124 | spl72_355), inference(avatar_contradiction_clause, [], [f7132])).
fof(f7132, plain, ($false | (~ spl72_73 | ~ spl72_124 | spl72_355)), inference(subsumption_resolution, [], [f7126, f1287])).
fof(f7126, plain, (~ (e23 = op2(e20, e21)) | (~ spl72_73 | spl72_355)), inference(backward_demodulation, [], [f2444, f1071])).
fof(f1071, plain, ((e20 = op2(e23, e21)) | ~ spl72_73), inference(avatar_component_clause, [], [f1069])).
fof(f2444, plain, (~ (e23 = op2(op2(e23, e21), e21)) | spl72_355), inference(avatar_component_clause, [], [f2442])).
fof(f2442, plain, (spl72_355 <=> (e23 = op2(op2(e23, e21), e21))), introduced(avatar_definition, [new_symbols(naming, [spl72_355])])).
fof(f7127, plain, (~ spl72_69 | ~ spl72_73), inference(avatar_split_clause, [], [f7118, f1069, f1052])).
fof(f7118, plain, (~ (e20 = op2(e23, e22)) | ~ spl72_73), inference(backward_demodulation, [], [f357, f1071])).
fof(f7112, plain, (spl72_86 | ~ spl72_91 | ~ spl72_370), inference(avatar_split_clause, [], [f7109, f2515, f1145, f1124])).
fof(f1124, plain, (spl72_86 <=> (e21 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl72_86])])).
fof(f2515, plain, (spl72_370 <=> (e21 = op2(op2(e22, e21), e22))), introduced(avatar_definition, [new_symbols(naming, [spl72_370])])).
fof(f7109, plain, ((e21 = op2(e22, e22)) | (~ spl72_91 | ~ spl72_370)), inference(backward_demodulation, [], [f2517, f1147])).
fof(f2517, plain, ((e21 = op2(op2(e22, e21), e22)) | ~ spl72_370), inference(avatar_component_clause, [], [f2515])).
fof(f7092, plain, (~ spl72_49 | ~ spl72_60 | spl72_204), inference(avatar_contradiction_clause, [], [f7091])).
fof(f7091, plain, ($false | (~ spl72_49 | ~ spl72_60 | spl72_204)), inference(subsumption_resolution, [], [f7090, f937])).
fof(f7090, plain, (~ (e10 = op1(e10, e13)) | (~ spl72_60 | spl72_204)), inference(forward_demodulation, [], [f1707, f983])).
fof(f1707, plain, (~ (e10 = op1(e10, op1(e10, e11))) | spl72_204), inference(avatar_component_clause, [], [f1705])).
fof(f1705, plain, (spl72_204 <=> (e10 = op1(e10, op1(e10, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl72_204])])).
fof(f7080, plain, (~ spl72_87 | ~ spl72_99 | spl72_311), inference(avatar_split_clause, [], [f7079, f2226, f1179, f1128])).
fof(f2226, plain, (spl72_311 <=> (e22 = op2(e22, op2(e21, e23)))), introduced(avatar_definition, [new_symbols(naming, [spl72_311])])).
fof(f7079, plain, (~ (e22 = op2(e22, e22)) | (~ spl72_99 | spl72_311)), inference(forward_demodulation, [], [f2228, f1181])).
fof(f2228, plain, (~ (e22 = op2(e22, op2(e21, e23))) | spl72_311), inference(avatar_component_clause, [], [f2226])).
fof(f7051, plain, (~ spl72_89 | ~ spl72_431), inference(avatar_split_clause, [], [f7040, f2865, f1137])).
fof(f7040, plain, (~ (e20 = op2(e22, e21)) | ~ spl72_431), inference(backward_demodulation, [], [f2740, f2866])).
fof(f2740, plain, ~ (op2(e22, e21) = h3(e12)), inference(backward_demodulation, [], [f351, f660])).
fof(f351, plain, ~ (op2(e22, e21) = op2(e22, e22)), inference(cnf_transformation, [], [f6])).
fof(f7050, plain, (~ spl72_69 | ~ spl72_431), inference(avatar_split_clause, [], [f7039, f2865, f1052])).
fof(f7039, plain, (~ (e20 = op2(e23, e22)) | ~ spl72_431), inference(backward_demodulation, [], [f2738, f2866])).
fof(f7013, plain, (~ spl72_427 | ~ spl72_118), inference(avatar_split_clause, [], [f3959, f1260, f2845])).
fof(f2845, plain, (spl72_427 <=> (e21 = h3(e12))), introduced(avatar_definition, [new_symbols(naming, [spl72_427])])).
fof(f3959, plain, (~ (e21 = h3(e12)) | ~ spl72_118), inference(backward_demodulation, [], [f2736, f1262])).
fof(f2736, plain, ~ (op2(e20, e22) = h3(e12)), inference(backward_demodulation, [], [f325, f660])).
fof(f325, plain, ~ (op2(e20, e22) = op2(e22, e22)), inference(cnf_transformation, [], [f6])).
fof(f6971, plain, (~ spl72_73 | ~ spl72_78 | spl72_361), inference(avatar_split_clause, [], [f6941, f2471, f1090, f1069])).
fof(f2471, plain, (spl72_361 <=> (e20 = op2(e23, op2(e23, e20)))), introduced(avatar_definition, [new_symbols(naming, [spl72_361])])).
fof(f6941, plain, (~ (e20 = op2(e23, e21)) | (~ spl72_78 | spl72_361)), inference(backward_demodulation, [], [f2473, f1092])).
fof(f2473, plain, (~ (e20 = op2(e23, op2(e23, e20))) | spl72_361), inference(avatar_component_clause, [], [f2471])).
fof(f6970, plain, (~ spl72_72 | ~ spl72_471), inference(avatar_split_clause, [], [f6969, f3076, f1064])).
fof(f6969, plain, (~ (e23 = op2(e23, e22)) | ~ spl72_471), inference(forward_demodulation, [], [f2750, f3077])).
fof(f2750, plain, ~ (op2(e23, e22) = h4(e12)), inference(backward_demodulation, [], [f359, f664])).
fof(f359, plain, ~ (op2(e23, e22) = op2(e23, e23)), inference(cnf_transformation, [], [f6])).
fof(f6963, plain, (spl72_69 | ~ spl72_96 | ~ spl72_374), inference(avatar_split_clause, [], [f6911, f2534, f1166, f1052])).
fof(f2534, plain, (spl72_374 <=> (e20 = op2(op2(e22, e20), e22))), introduced(avatar_definition, [new_symbols(naming, [spl72_374])])).
fof(f6911, plain, ((e20 = op2(e23, e22)) | (~ spl72_96 | ~ spl72_374)), inference(backward_demodulation, [], [f2536, f1168])).
fof(f2536, plain, ((e20 = op2(op2(e22, e20), e22)) | ~ spl72_374), inference(avatar_component_clause, [], [f2534])).
fof(f6917, plain, (~ spl72_96 | ~ spl72_113 | spl72_307), inference(avatar_contradiction_clause, [], [f6916])).
fof(f6916, plain, ($false | (~ spl72_96 | ~ spl72_113 | spl72_307)), inference(subsumption_resolution, [], [f6909, f1241])).
fof(f6909, plain, (~ (e20 = op2(e20, e23)) | (~ spl72_96 | spl72_307)), inference(backward_demodulation, [], [f2209, f1168])).
fof(f2209, plain, (~ (e20 = op2(e20, op2(e22, e20))) | spl72_307), inference(avatar_component_clause, [], [f2207])).
fof(f2207, plain, (spl72_307 <=> (e20 = op2(e20, op2(e22, e20)))), introduced(avatar_definition, [new_symbols(naming, [spl72_307])])).
fof(f6915, plain, (~ spl72_68 | ~ spl72_96 | spl72_304), inference(avatar_split_clause, [], [f6906, f2192, f1166, f1047])).
fof(f6906, plain, (~ (e23 = op2(e23, e23)) | (~ spl72_96 | spl72_304)), inference(backward_demodulation, [], [f2194, f1168])).
fof(f2194, plain, (~ (e23 = op2(e23, op2(e22, e20))) | spl72_304), inference(avatar_component_clause, [], [f2192])).
fof(f6914, plain, (~ spl72_84 | ~ spl72_96), inference(avatar_split_clause, [], [f6905, f1166, f1115])).
fof(f6905, plain, (~ (e23 = op2(e22, e23)) | ~ spl72_96), inference(backward_demodulation, [], [f350, f1168])).
fof(f350, plain, ~ (op2(e22, e20) = op2(e22, e23)), inference(cnf_transformation, [], [f6])).
fof(f6891, plain, (~ spl72_32 | ~ spl72_54 | ~ spl72_168), inference(avatar_contradiction_clause, [], [f6890])).
fof(f6890, plain, ($false | (~ spl72_32 | ~ spl72_54 | ~ spl72_168)), inference(subsumption_resolution, [], [f6889, f364])).
fof(f6889, plain, ((e11 = e13) | (~ spl72_32 | ~ spl72_54 | ~ spl72_168)), inference(forward_demodulation, [], [f6888, f958])).
fof(f6888, plain, ((e13 = op1(e10, e12)) | (~ spl72_32 | ~ spl72_168)), inference(forward_demodulation, [], [f1530, f864])).
fof(f1530, plain, ((op1(e10, e12) = op1(e12, e10)) | ~ spl72_168), inference(avatar_component_clause, [], [f1528])).
fof(f1528, plain, (spl72_168 <=> (op1(e10, e12) = op1(e12, e10))), introduced(avatar_definition, [new_symbols(naming, [spl72_168])])).
fof(f6874, plain, (spl72_75 | ~ spl72_104 | ~ spl72_382), inference(avatar_split_clause, [], [f6873, f2572, f1200, f1077])).
fof(f2572, plain, (spl72_382 <=> (e22 = op2(op2(e21, e22), e21))), introduced(avatar_definition, [new_symbols(naming, [spl72_382])])).
fof(f6873, plain, ((e22 = op2(e23, e21)) | (~ spl72_104 | ~ spl72_382)), inference(forward_demodulation, [], [f2574, f1202])).
fof(f2574, plain, ((e22 = op2(op2(e21, e22), e21)) | ~ spl72_382), inference(avatar_component_clause, [], [f2572])).
fof(f6845, plain, (~ spl72_84 | ~ spl72_471), inference(avatar_split_clause, [], [f6839, f3076, f1115])).
fof(f6839, plain, (~ (e23 = op2(e22, e23)) | ~ spl72_471), inference(backward_demodulation, [], [f2747, f3077])).
fof(f6804, plain, (spl72_94 | ~ spl72_118 | ~ spl72_308), inference(avatar_split_clause, [], [f6803, f2212, f1260, f1158])).
fof(f2212, plain, (spl72_308 <=> (op2(e20, e22) = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl72_308])])).
fof(f6803, plain, ((e21 = op2(e22, e20)) | (~ spl72_118 | ~ spl72_308)), inference(forward_demodulation, [], [f2214, f1262])).
fof(f2214, plain, ((op2(e20, e22) = op2(e22, e20)) | ~ spl72_308), inference(avatar_component_clause, [], [f2212])).
fof(f6799, plain, (spl72_94 | ~ spl72_89 | ~ spl72_371), inference(avatar_split_clause, [], [f6722, f2520, f1137, f1158])).
fof(f2520, plain, (spl72_371 <=> (e21 = op2(e22, op2(e22, e21)))), introduced(avatar_definition, [new_symbols(naming, [spl72_371])])).
fof(f6722, plain, ((e21 = op2(e22, e20)) | (~ spl72_89 | ~ spl72_371)), inference(forward_demodulation, [], [f2521, f1139])).
fof(f1139, plain, ((e20 = op2(e22, e21)) | ~ spl72_89), inference(avatar_component_clause, [], [f1137])).
fof(f2521, plain, ((e21 = op2(e22, op2(e22, e21))) | ~ spl72_371), inference(avatar_component_clause, [], [f2520])).
fof(f6771, plain, (~ spl72_98 | spl72_377), inference(avatar_contradiction_clause, [], [f6770])).
fof(f6770, plain, ($false | (~ spl72_98 | spl72_377)), inference(subsumption_resolution, [], [f6764, f1177])).
fof(f6764, plain, (~ (e21 = op2(e21, e23)) | (~ spl72_98 | spl72_377)), inference(backward_demodulation, [], [f2550, f1177])).
fof(f2550, plain, (~ (e21 = op2(op2(e21, e23), e23)) | spl72_377), inference(avatar_component_clause, [], [f2548])).
fof(f2548, plain, (spl72_377 <=> (e21 = op2(op2(e21, e23), e23))), introduced(avatar_definition, [new_symbols(naming, [spl72_377])])).
fof(f6713, plain, (~ spl72_75 | ~ spl72_411), inference(avatar_split_clause, [], [f6705, f2764, f1077])).
fof(f6705, plain, (~ (e22 = op2(e23, e21)) | ~ spl72_411), inference(backward_demodulation, [], [f2749, f2765])).
fof(f2749, plain, ~ (op2(e23, e21) = h4(e12)), inference(backward_demodulation, [], [f358, f664])).
fof(f358, plain, ~ (op2(e23, e21) = op2(e23, e23)), inference(cnf_transformation, [], [f6])).
fof(f6582, plain, (~ spl72_23 | spl72_154), inference(avatar_contradiction_clause, [], [f6581])).
fof(f6581, plain, ($false | (~ spl72_23 | spl72_154)), inference(subsumption_resolution, [], [f6577, f826])).
fof(f6577, plain, (~ (e12 = op1(e12, e12)) | (~ spl72_23 | spl72_154)), inference(backward_demodulation, [], [f1462, f826])).
fof(f1462, plain, (~ (e12 = op1(e12, op1(e12, e12))) | spl72_154), inference(avatar_component_clause, [], [f1460])).
fof(f1460, plain, (spl72_154 <=> (e12 = op1(e12, op1(e12, e12)))), introduced(avatar_definition, [new_symbols(naming, [spl72_154])])).
fof(f6558, plain, (~ spl72_36 | spl72_239), inference(avatar_contradiction_clause, [], [f6557])).
fof(f6557, plain, ($false | (~ spl72_36 | spl72_239)), inference(subsumption_resolution, [], [f6545, f881])).
fof(f881, plain, ((e13 = op1(e11, e13)) | ~ spl72_36), inference(avatar_component_clause, [], [f879])).
fof(f6545, plain, (~ (e13 = op1(e11, e13)) | (~ spl72_36 | spl72_239)), inference(backward_demodulation, [], [f1876, f881])).
fof(f6532, plain, (~ spl72_37 | ~ spl72_46 | spl72_177), inference(avatar_contradiction_clause, [], [f6531])).
fof(f6531, plain, ($false | (~ spl72_37 | ~ spl72_46 | spl72_177)), inference(subsumption_resolution, [], [f6516, f924])).
fof(f6516, plain, (~ (e11 = op1(e11, e10)) | (~ spl72_37 | spl72_177)), inference(backward_demodulation, [], [f1574, f886])).
fof(f1574, plain, (~ (e11 = op1(e11, op1(e11, e12))) | spl72_177), inference(avatar_component_clause, [], [f1572])).
fof(f1572, plain, (spl72_177 <=> (e11 = op1(e11, op1(e11, e12)))), introduced(avatar_definition, [new_symbols(naming, [spl72_177])])).
fof(f6528, plain, (~ spl72_21 | ~ spl72_37), inference(avatar_split_clause, [], [f6508, f884, f816])).
fof(f6508, plain, (~ (e10 = op1(e12, e12)) | ~ spl72_37), inference(backward_demodulation, [], [f279, f886])).
fof(f6504, plain, (~ spl72_39 | ~ spl72_43), inference(avatar_split_clause, [], [f6498, f909, f892])).
fof(f6498, plain, (~ (e12 = op1(e11, e12)) | ~ spl72_43), inference(backward_demodulation, [], [f297, f911])).
fof(f6503, plain, (~ spl72_27 | ~ spl72_43), inference(avatar_split_clause, [], [f6497, f909, f841])).
fof(f6497, plain, (~ (e12 = op1(e12, e11)) | ~ spl72_43), inference(backward_demodulation, [], [f273, f911])).
fof(f6495, plain, (spl72_41 | ~ spl72_46 | ~ spl72_248), inference(avatar_split_clause, [], [f6488, f1918, f922, f901])).
fof(f1918, plain, (spl72_248 <=> (e10 = op1(op1(e11, e10), e11))), introduced(avatar_definition, [new_symbols(naming, [spl72_248])])).
fof(f6488, plain, ((e10 = op1(e11, e11)) | (~ spl72_46 | ~ spl72_248)), inference(backward_demodulation, [], [f1920, f924])).
fof(f1920, plain, ((e10 = op1(op1(e11, e10), e11)) | ~ spl72_248), inference(avatar_component_clause, [], [f1918])).
fof(f6494, plain, (~ spl72_46 | spl72_247), inference(avatar_contradiction_clause, [], [f6493])).
fof(f6493, plain, ($false | (~ spl72_46 | spl72_247)), inference(subsumption_resolution, [], [f6487, f924])).
fof(f6487, plain, (~ (e11 = op1(e11, e10)) | (~ spl72_46 | spl72_247)), inference(backward_demodulation, [], [f1915, f924])).
fof(f6492, plain, (~ spl72_34 | ~ spl72_46), inference(avatar_split_clause, [], [f6484, f922, f871])).
fof(f6484, plain, (~ (e11 = op1(e11, e13)) | ~ spl72_46), inference(backward_demodulation, [], [f296, f924])).
fof(f6491, plain, (~ spl72_30 | ~ spl72_46), inference(avatar_split_clause, [], [f6481, f922, f854])).
fof(f6481, plain, (~ (e11 = op1(e12, e10)) | ~ spl72_46), inference(backward_demodulation, [], [f267, f924])).
fof(f6462, plain, (~ spl72_89 | ~ spl72_94 | spl72_371), inference(avatar_contradiction_clause, [], [f6461])).
fof(f6461, plain, ($false | (~ spl72_89 | ~ spl72_94 | spl72_371)), inference(subsumption_resolution, [], [f6457, f1160])).
fof(f1160, plain, ((e21 = op2(e22, e20)) | ~ spl72_94), inference(avatar_component_clause, [], [f1158])).
fof(f6457, plain, (~ (e21 = op2(e22, e20)) | (~ spl72_89 | spl72_371)), inference(backward_demodulation, [], [f2522, f1139])).
fof(f2522, plain, (~ (e21 = op2(e22, op2(e22, e21))) | spl72_371), inference(avatar_component_clause, [], [f2520])).
fof(f6396, plain, (~ spl72_71 | ~ spl72_423), inference(avatar_split_clause, [], [f6381, f2824, f1060])).
fof(f6381, plain, (~ (e22 = op2(e23, e22)) | ~ spl72_423), inference(backward_demodulation, [], [f2738, f2825])).
fof(f6347, plain, (~ spl72_23 | ~ spl72_3 | spl72_265), inference(avatar_split_clause, [], [f6327, f2001, f739, f824])).
fof(f739, plain, (spl72_3 <=> (e12 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl72_3])])).
fof(f2001, plain, (spl72_265 <=> (e12 = op1(e12, op1(e13, e13)))), introduced(avatar_definition, [new_symbols(naming, [spl72_265])])).
fof(f6327, plain, (~ (e12 = op1(e12, e12)) | (~ spl72_3 | spl72_265)), inference(backward_demodulation, [], [f2003, f741])).
fof(f741, plain, ((e12 = op1(e13, e13)) | ~ spl72_3), inference(avatar_component_clause, [], [f739])).
fof(f2003, plain, (~ (e12 = op1(e12, op1(e13, e13))) | spl72_265), inference(avatar_component_clause, [], [f2001])).
fof(f6343, plain, (~ spl72_18 | ~ spl72_10 | ~ spl72_134), inference(avatar_split_clause, [], [f6342, f1363, f769, f803])).
fof(f1363, plain, (spl72_134 <=> (op1(e12, e13) = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl72_134])])).
fof(f6342, plain, (~ (e11 = op1(e12, e13)) | (~ spl72_10 | ~ spl72_134)), inference(forward_demodulation, [], [f5292, f771])).
fof(f5292, plain, (~ (op1(e12, e13) = op1(e13, e11)) | ~ spl72_134), inference(forward_demodulation, [], [f309, f1365])).
fof(f1365, plain, ((op1(e12, e13) = op1(e13, e12)) | ~ spl72_134), inference(avatar_component_clause, [], [f1363])).
fof(f6332, plain, (~ spl72_3 | ~ spl72_8 | spl72_264), inference(avatar_contradiction_clause, [], [f6331])).
fof(f6331, plain, ($false | (~ spl72_3 | ~ spl72_8 | spl72_264)), inference(subsumption_resolution, [], [f6326, f762])).
fof(f762, plain, ((e13 = op1(e13, e12)) | ~ spl72_8), inference(avatar_component_clause, [], [f760])).
fof(f6326, plain, (~ (e13 = op1(e13, e12)) | (~ spl72_3 | spl72_264)), inference(backward_demodulation, [], [f1998, f741])).
fof(f1998, plain, (~ (e13 = op1(e13, op1(e13, e13))) | spl72_264), inference(avatar_component_clause, [], [f1996])).
fof(f1996, plain, (spl72_264 <=> (e13 = op1(e13, op1(e13, e13)))), introduced(avatar_definition, [new_symbols(naming, [spl72_264])])).
fof(f6318, plain, (~ spl72_13 | ~ spl72_21 | ~ spl72_153), inference(avatar_contradiction_clause, [], [f6317])).
fof(f6317, plain, ($false | (~ spl72_13 | ~ spl72_21 | ~ spl72_153)), inference(subsumption_resolution, [], [f6316, f362])).
fof(f362, plain, ~ (e10 = e13), inference(cnf_transformation, [], [f7])).
fof(f6316, plain, ((e10 = e13) | (~ spl72_13 | ~ spl72_21 | ~ spl72_153)), inference(forward_demodulation, [], [f6311, f784])).
fof(f6311, plain, ((e13 = op1(e13, e10)) | (~ spl72_21 | ~ spl72_153)), inference(backward_demodulation, [], [f1456, f818])).
fof(f818, plain, ((e10 = op1(e12, e12)) | ~ spl72_21), inference(avatar_component_clause, [], [f816])).
fof(f1456, plain, ((e13 = op1(e13, op1(e12, e12))) | ~ spl72_153), inference(avatar_component_clause, [], [f1455])).
fof(f1455, plain, (spl72_153 <=> (e13 = op1(e13, op1(e12, e12)))), introduced(avatar_definition, [new_symbols(naming, [spl72_153])])).
fof(f6302, plain, (~ spl72_34 | spl72_237), inference(avatar_contradiction_clause, [], [f6301])).
fof(f6301, plain, ($false | (~ spl72_34 | spl72_237)), inference(subsumption_resolution, [], [f6292, f873])).
fof(f6292, plain, (~ (e11 = op1(e11, e13)) | (~ spl72_34 | spl72_237)), inference(backward_demodulation, [], [f1866, f873])).
fof(f6275, plain, (~ spl72_39 | spl72_243), inference(avatar_contradiction_clause, [], [f6274])).
fof(f6274, plain, ($false | (~ spl72_39 | spl72_243)), inference(subsumption_resolution, [], [f6263, f894])).
fof(f6263, plain, (~ (e12 = op1(e11, e12)) | (~ spl72_39 | spl72_243)), inference(backward_demodulation, [], [f1895, f894])).
fof(f6204, plain, (spl72_3 | ~ spl72_8 | ~ spl72_212), inference(avatar_split_clause, [], [f6203, f1744, f760, f739])).
fof(f1744, plain, (spl72_212 <=> (e12 = op1(op1(e13, e12), e13))), introduced(avatar_definition, [new_symbols(naming, [spl72_212])])).
fof(f6203, plain, ((e12 = op1(e13, e13)) | (~ spl72_8 | ~ spl72_212)), inference(forward_demodulation, [], [f1746, f762])).
fof(f1746, plain, ((e12 = op1(op1(e13, e12), e13)) | ~ spl72_212), inference(avatar_component_clause, [], [f1744])).
fof(f6188, plain, (~ spl72_106 | spl72_322), inference(avatar_contradiction_clause, [], [f6187])).
fof(f6187, plain, ($false | (~ spl72_106 | spl72_322)), inference(subsumption_resolution, [], [f6186, f1211])).
fof(f1211, plain, ((e21 = op2(e21, e21)) | ~ spl72_106), inference(avatar_component_clause, [], [f1209])).
fof(f1209, plain, (spl72_106 <=> (e21 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl72_106])])).
fof(f6186, plain, (~ (e21 = op2(e21, e21)) | (~ spl72_106 | spl72_322)), inference(forward_demodulation, [], [f2283, f1211])).
fof(f2283, plain, (~ (e21 = op2(e21, op2(e21, e21))) | spl72_322), inference(avatar_component_clause, [], [f2281])).
fof(f2281, plain, (spl72_322 <=> (e21 = op2(e21, op2(e21, e21)))), introduced(avatar_definition, [new_symbols(naming, [spl72_322])])).
fof(f6182, plain, (~ spl72_94 | spl72_101 | ~ spl72_374), inference(avatar_contradiction_clause, [], [f6181])).
fof(f6181, plain, ($false | (~ spl72_94 | spl72_101 | ~ spl72_374)), inference(subsumption_resolution, [], [f6180, f1189])).
fof(f1189, plain, (~ (e20 = op2(e21, e22)) | spl72_101), inference(avatar_component_clause, [], [f1188])).
fof(f6180, plain, ((e20 = op2(e21, e22)) | (~ spl72_94 | ~ spl72_374)), inference(forward_demodulation, [], [f2536, f1160])).
fof(f6179, plain, (~ spl72_89 | ~ spl72_94 | spl72_375), inference(avatar_split_clause, [], [f6178, f2539, f1158, f1137])).
fof(f2539, plain, (spl72_375 <=> (e20 = op2(e22, op2(e22, e20)))), introduced(avatar_definition, [new_symbols(naming, [spl72_375])])).
fof(f6178, plain, (~ (e20 = op2(e22, e21)) | (~ spl72_94 | spl72_375)), inference(forward_demodulation, [], [f2541, f1160])).
fof(f2541, plain, (~ (e20 = op2(e22, op2(e22, e20))) | spl72_375), inference(avatar_component_clause, [], [f2539])).
fof(f6171, plain, (spl72_91 | ~ spl72_66 | ~ spl72_405), inference(avatar_split_clause, [], [f6170, f2685, f1039, f1145])).
fof(f1039, plain, (spl72_66 <=> (e21 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl72_66])])).
fof(f6170, plain, ((e22 = op2(e22, e21)) | (~ spl72_66 | ~ spl72_405)), inference(forward_demodulation, [], [f2686, f1041])).
fof(f1041, plain, ((e21 = op2(e23, e23)) | ~ spl72_66), inference(avatar_component_clause, [], [f1039])).
fof(f2686, plain, ((e22 = op2(e22, op2(e23, e23))) | ~ spl72_405), inference(avatar_component_clause, [], [f2685])).
fof(f6101, plain, (spl72_37 | ~ spl72_30 | ~ spl72_234), inference(avatar_split_clause, [], [f5577, f1850, f854, f884])).
fof(f5577, plain, ((e10 = op1(e11, e12)) | (~ spl72_30 | ~ spl72_234)), inference(forward_demodulation, [], [f1852, f856])).
fof(f6081, plain, (spl72_241 | ~ spl72_162 | ~ spl72_230), inference(avatar_split_clause, [], [f6080, f1831, f1499, f1883])).
fof(f1831, plain, (spl72_230 <=> (e11 = op1(op1(e12, e11), e12))), introduced(avatar_definition, [new_symbols(naming, [spl72_230])])).
fof(f6080, plain, ((e11 = op1(op1(e11, e12), e12)) | (~ spl72_162 | ~ spl72_230)), inference(forward_demodulation, [], [f1833, f1501])).
fof(f1501, plain, ((op1(e11, e12) = op1(e12, e11)) | ~ spl72_162), inference(avatar_component_clause, [], [f1499])).
fof(f1833, plain, ((e11 = op1(op1(e12, e11), e12)) | ~ spl72_230), inference(avatar_component_clause, [], [f1831])).
fof(f6076, plain, (~ spl72_177 | spl72_160 | ~ spl72_162), inference(avatar_split_clause, [], [f6071, f1499, f1489, f1572])).
fof(f1489, plain, (spl72_160 <=> (e11 = op1(e11, op1(e12, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl72_160])])).
fof(f6071, plain, (~ (e11 = op1(e11, op1(e11, e12))) | (spl72_160 | ~ spl72_162)), inference(backward_demodulation, [], [f1491, f1501])).
fof(f1491, plain, (~ (e11 = op1(e11, op1(e12, e11))) | spl72_160), inference(avatar_component_clause, [], [f1489])).
fof(f6059, plain, (~ spl72_37 | ~ spl72_54 | spl72_241), inference(avatar_contradiction_clause, [], [f6058])).
fof(f6058, plain, ($false | (~ spl72_37 | ~ spl72_54 | spl72_241)), inference(subsumption_resolution, [], [f6057, f958])).
fof(f6057, plain, (~ (e11 = op1(e10, e12)) | (~ spl72_37 | spl72_241)), inference(forward_demodulation, [], [f1885, f886])).
fof(f1885, plain, (~ (e11 = op1(op1(e11, e12), e12)) | spl72_241), inference(avatar_component_clause, [], [f1883])).
fof(f5966, plain, (~ spl72_77 | ~ spl72_109), inference(avatar_split_clause, [], [f5956, f1222, f1086])).
fof(f5956, plain, (~ (e20 = op2(e23, e20)) | ~ spl72_109), inference(backward_demodulation, [], [f316, f1224])).
fof(f1224, plain, ((e20 = op2(e21, e20)) | ~ spl72_109), inference(avatar_component_clause, [], [f1222])).
fof(f5953, plain, (~ spl72_49 | spl72_251), inference(avatar_contradiction_clause, [], [f5952])).
fof(f5952, plain, ($false | (~ spl72_49 | spl72_251)), inference(subsumption_resolution, [], [f5951, f937])).
fof(f5951, plain, (~ (e10 = op1(e10, e13)) | (~ spl72_49 | spl72_251)), inference(forward_demodulation, [], [f1934, f937])).
fof(f1934, plain, (~ (e10 = op1(op1(e10, e13), e13)) | spl72_251), inference(avatar_component_clause, [], [f1932])).
fof(f1932, plain, (spl72_251 <=> (e10 = op1(op1(e10, e13), e13))), introduced(avatar_definition, [new_symbols(naming, [spl72_251])])).
fof(f5856, plain, (~ spl72_77 | spl72_361), inference(avatar_contradiction_clause, [], [f5855])).
fof(f5855, plain, ($false | (~ spl72_77 | spl72_361)), inference(subsumption_resolution, [], [f5843, f1088])).
fof(f5843, plain, (~ (e20 = op2(e23, e20)) | (~ spl72_77 | spl72_361)), inference(backward_demodulation, [], [f2473, f1088])).
fof(f5852, plain, (~ spl72_77 | ~ spl72_127 | ~ spl72_359), inference(avatar_contradiction_clause, [], [f5851])).
fof(f5851, plain, ($false | (~ spl72_77 | ~ spl72_127 | ~ spl72_359)), inference(subsumption_resolution, [], [f5850, f371])).
fof(f5850, plain, ((e22 = e23) | (~ spl72_77 | ~ spl72_127 | ~ spl72_359)), inference(forward_demodulation, [], [f5841, f1300])).
fof(f5841, plain, ((op2(e20, e20) = e23) | (~ spl72_77 | ~ spl72_359)), inference(backward_demodulation, [], [f2462, f1088])).
fof(f2462, plain, ((e23 = op2(op2(e23, e20), e20)) | ~ spl72_359), inference(avatar_component_clause, [], [f2461])).
fof(f5845, plain, (~ spl72_69 | ~ spl72_77), inference(avatar_split_clause, [], [f5833, f1086, f1052])).
fof(f5833, plain, (~ (e20 = op2(e23, e22)) | ~ spl72_77), inference(backward_demodulation, [], [f355, f1088])).
fof(f355, plain, ~ (op2(e23, e20) = op2(e23, e22)), inference(cnf_transformation, [], [f6])).
fof(f5826, plain, (~ spl72_91 | spl72_369), inference(avatar_contradiction_clause, [], [f5825])).
fof(f5825, plain, ($false | (~ spl72_91 | spl72_369)), inference(subsumption_resolution, [], [f5821, f1147])).
fof(f5821, plain, (~ (e22 = op2(e22, e21)) | (~ spl72_91 | spl72_369)), inference(backward_demodulation, [], [f2512, f1147])).
fof(f2512, plain, (~ (e22 = op2(op2(e22, e21), e21)) | spl72_369), inference(avatar_component_clause, [], [f2510])).
fof(f2510, plain, (spl72_369 <=> (e22 = op2(op2(e22, e21), e21))), introduced(avatar_definition, [new_symbols(naming, [spl72_369])])).
fof(f5810, plain, (spl72_73 | ~ spl72_112 | ~ spl72_388), inference(avatar_split_clause, [], [f5804, f2602, f1234, f1069])).
fof(f2602, plain, (spl72_388 <=> (e20 = op2(op2(e21, e20), e21))), introduced(avatar_definition, [new_symbols(naming, [spl72_388])])).
fof(f5804, plain, ((e20 = op2(e23, e21)) | (~ spl72_112 | ~ spl72_388)), inference(backward_demodulation, [], [f2604, f1236])).
fof(f2604, plain, ((e20 = op2(op2(e21, e20), e21)) | ~ spl72_388), inference(avatar_component_clause, [], [f2602])).
fof(f5809, plain, (~ spl72_112 | ~ spl72_113 | spl72_328), inference(avatar_contradiction_clause, [], [f5808])).
fof(f5808, plain, ($false | (~ spl72_112 | ~ spl72_113 | spl72_328)), inference(subsumption_resolution, [], [f5802, f1241])).
fof(f5802, plain, (~ (e20 = op2(e20, e23)) | (~ spl72_112 | spl72_328)), inference(backward_demodulation, [], [f2312, f1236])).
fof(f2312, plain, (~ (e20 = op2(e20, op2(e21, e20))) | spl72_328), inference(avatar_component_clause, [], [f2310])).
fof(f5796, plain, (~ spl72_25 | ~ spl72_35 | ~ spl72_238), inference(avatar_contradiction_clause, [], [f5795])).
fof(f5795, plain, ($false | (~ spl72_25 | ~ spl72_35 | ~ spl72_238)), inference(subsumption_resolution, [], [f5794, f362])).
fof(f5794, plain, ((e10 = e13) | (~ spl72_25 | ~ spl72_35 | ~ spl72_238)), inference(forward_demodulation, [], [f5793, f835])).
fof(f5793, plain, ((e13 = op1(e12, e11)) | (~ spl72_35 | ~ spl72_238)), inference(forward_demodulation, [], [f1871, f877])).
fof(f5726, plain, (~ spl72_97 | ~ spl72_113), inference(avatar_split_clause, [], [f4468, f1239, f1171])).
fof(f1171, plain, (spl72_97 <=> (e20 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl72_97])])).
fof(f4468, plain, (~ (e20 = op2(e21, e23)) | ~ spl72_113), inference(forward_demodulation, [], [f330, f1241])).
fof(f330, plain, ~ (op2(e20, e23) = op2(e21, e23)), inference(cnf_transformation, [], [f6])).
fof(f5710, plain, (~ spl72_76 | ~ spl72_515), inference(avatar_split_clause, [], [f3932, f3362, f1081])).
fof(f1081, plain, (spl72_76 <=> (e23 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl72_76])])).
fof(f3932, plain, (~ (e23 = op2(e23, e21)) | ~ spl72_515), inference(forward_demodulation, [], [f2715, f3363])).
fof(f2715, plain, ~ (op2(e23, e21) = h1(e13)), inference(backward_demodulation, [], [f320, f2712])).
fof(f2712, plain, (op2(e20, e21) = h1(e13)), inference(forward_demodulation, [], [f653, f647])).
fof(f653, plain, (op2(e20, op2(e20, op2(e20, e20))) = h1(e13)), inference(cnf_transformation, [], [f14])).
fof(f320, plain, ~ (op2(e20, e21) = op2(e23, e21)), inference(cnf_transformation, [], [f6])).
fof(f5659, plain, (~ spl72_109 | spl72_389), inference(avatar_contradiction_clause, [], [f5658])).
fof(f5658, plain, ($false | (~ spl72_109 | spl72_389)), inference(subsumption_resolution, [], [f5650, f1224])).
fof(f5650, plain, (~ (e20 = op2(e21, e20)) | (~ spl72_109 | spl72_389)), inference(backward_demodulation, [], [f2609, f1224])).
fof(f2609, plain, (~ (e20 = op2(e21, op2(e21, e20))) | spl72_389), inference(avatar_component_clause, [], [f2607])).
fof(f2607, plain, (spl72_389 <=> (e20 = op2(e21, op2(e21, e20)))), introduced(avatar_definition, [new_symbols(naming, [spl72_389])])).
fof(f5635, plain, (spl72_47 | ~ spl72_54 | ~ spl72_256), inference(avatar_contradiction_clause, [], [f5634])).
fof(f5634, plain, ($false | (spl72_47 | ~ spl72_54 | ~ spl72_256)), inference(subsumption_resolution, [], [f5633, f927])).
fof(f927, plain, (~ (e12 = op1(e11, e10)) | spl72_47), inference(avatar_component_clause, [], [f926])).
fof(f926, plain, (spl72_47 <=> (e12 = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl72_47])])).
fof(f5633, plain, ((e12 = op1(e11, e10)) | (~ spl72_54 | ~ spl72_256)), inference(forward_demodulation, [], [f1958, f958])).
fof(f1958, plain, ((e12 = op1(op1(e10, e12), e10)) | ~ spl72_256), inference(avatar_component_clause, [], [f1956])).
fof(f1956, plain, (spl72_256 <=> (e12 = op1(op1(e10, e12), e10))), introduced(avatar_definition, [new_symbols(naming, [spl72_256])])).
fof(f5631, plain, (spl72_92 | ~ spl72_104 | ~ spl72_302), inference(avatar_contradiction_clause, [], [f5630])).
fof(f5630, plain, ($false | (spl72_92 | ~ spl72_104 | ~ spl72_302)), inference(subsumption_resolution, [], [f5629, f1150])).
fof(f5629, plain, ((e23 = op2(e22, e21)) | (~ spl72_104 | ~ spl72_302)), inference(forward_demodulation, [], [f2185, f1202])).
fof(f5576, plain, (~ spl72_25 | ~ spl72_30 | spl72_235), inference(avatar_contradiction_clause, [], [f5575])).
fof(f5575, plain, ($false | (~ spl72_25 | ~ spl72_30 | spl72_235)), inference(subsumption_resolution, [], [f5574, f835])).
fof(f5574, plain, (~ (e10 = op1(e12, e11)) | (~ spl72_30 | spl72_235)), inference(forward_demodulation, [], [f1857, f856])).
fof(f5573, plain, (~ spl72_110 | spl72_387), inference(avatar_contradiction_clause, [], [f5572])).
fof(f5572, plain, ($false | (~ spl72_110 | spl72_387)), inference(subsumption_resolution, [], [f5571, f1228])).
fof(f5571, plain, (~ (e21 = op2(e21, e20)) | (~ spl72_110 | spl72_387)), inference(forward_demodulation, [], [f2599, f1228])).
fof(f5542, plain, (~ spl72_25 | ~ spl72_30 | spl72_231), inference(avatar_contradiction_clause, [], [f5541])).
fof(f5541, plain, ($false | (~ spl72_25 | ~ spl72_30 | spl72_231)), inference(subsumption_resolution, [], [f5540, f856])).
fof(f5540, plain, (~ (e11 = op1(e12, e10)) | (~ spl72_25 | spl72_231)), inference(forward_demodulation, [], [f1838, f835])).
fof(f1838, plain, (~ (e11 = op1(e12, op1(e12, e11))) | spl72_231), inference(avatar_component_clause, [], [f1836])).
fof(f1836, plain, (spl72_231 <=> (e11 = op1(e12, op1(e12, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl72_231])])).
fof(f5533, plain, (~ spl72_113 | spl72_391), inference(avatar_contradiction_clause, [], [f5532])).
fof(f5532, plain, ($false | (~ spl72_113 | spl72_391)), inference(subsumption_resolution, [], [f5531, f1241])).
fof(f5531, plain, (~ (e20 = op2(e20, e23)) | (~ spl72_113 | spl72_391)), inference(forward_demodulation, [], [f2618, f1241])).
fof(f2618, plain, (~ (e20 = op2(op2(e20, e23), e23)) | spl72_391), inference(avatar_component_clause, [], [f2616])).
fof(f2616, plain, (spl72_391 <=> (e20 = op2(op2(e20, e23), e23))), introduced(avatar_definition, [new_symbols(naming, [spl72_391])])).
fof(f5511, plain, (~ spl72_57 | ~ spl72_60), inference(avatar_contradiction_clause, [], [f5510])).
fof(f5510, plain, ($false | (~ spl72_57 | ~ spl72_60)), inference(subsumption_resolution, [], [f5509, f362])).
fof(f5509, plain, ((e10 = e13) | (~ spl72_57 | ~ spl72_60)), inference(backward_demodulation, [], [f983, f971])).
fof(f971, plain, ((e10 = op1(e10, e11)) | ~ spl72_57), inference(avatar_component_clause, [], [f969])).
fof(f969, plain, (spl72_57 <=> (e10 = op1(e10, e11))), introduced(avatar_definition, [new_symbols(naming, [spl72_57])])).
fof(f5495, plain, (spl72_9 | ~ spl72_48 | ~ spl72_248), inference(avatar_contradiction_clause, [], [f5494])).
fof(f5494, plain, ($false | (spl72_9 | ~ spl72_48 | ~ spl72_248)), inference(subsumption_resolution, [], [f5493, f766])).
fof(f766, plain, (~ (e10 = op1(e13, e11)) | spl72_9), inference(avatar_component_clause, [], [f765])).
fof(f5493, plain, ((e10 = op1(e13, e11)) | (~ spl72_48 | ~ spl72_248)), inference(forward_demodulation, [], [f1920, f932])).
fof(f5413, plain, (~ spl72_13 | ~ spl72_49 | spl72_220), inference(avatar_contradiction_clause, [], [f5412])).
fof(f5412, plain, ($false | (~ spl72_13 | ~ spl72_49 | spl72_220)), inference(subsumption_resolution, [], [f5409, f937])).
fof(f5409, plain, (~ (e10 = op1(e10, e13)) | (~ spl72_13 | spl72_220)), inference(backward_demodulation, [], [f1783, f784])).
fof(f1783, plain, (~ (e10 = op1(op1(e13, e10), e13)) | spl72_220), inference(avatar_component_clause, [], [f1782])).
fof(f1782, plain, (spl72_220 <=> (e10 = op1(op1(e13, e10), e13))), introduced(avatar_definition, [new_symbols(naming, [spl72_220])])).
fof(f5397, plain, (~ spl72_20 | spl72_225), inference(avatar_contradiction_clause, [], [f5396])).
fof(f5396, plain, ($false | (~ spl72_20 | spl72_225)), inference(subsumption_resolution, [], [f5387, f813])).
fof(f813, plain, ((e13 = op1(e12, e13)) | ~ spl72_20), inference(avatar_component_clause, [], [f811])).
fof(f5387, plain, (~ (e13 = op1(e12, e13)) | (~ spl72_20 | spl72_225)), inference(backward_demodulation, [], [f1808, f813])).
fof(f1808, plain, (~ (e13 = op1(e12, op1(e12, e13))) | spl72_225), inference(avatar_component_clause, [], [f1806])).
fof(f1806, plain, (spl72_225 <=> (e13 = op1(e12, op1(e12, e13)))), introduced(avatar_definition, [new_symbols(naming, [spl72_225])])).
fof(f5381, plain, (~ spl72_14 | ~ spl72_30), inference(avatar_split_clause, [], [f5378, f854, f786])).
fof(f5378, plain, (~ (e11 = op1(e13, e10)) | ~ spl72_30), inference(backward_demodulation, [], [f269, f856])).
fof(f269, plain, ~ (op1(e12, e10) = op1(e13, e10)), inference(cnf_transformation, [], [f5])).
fof(f5330, plain, (~ spl72_99 | ~ spl72_104 | spl72_383), inference(avatar_contradiction_clause, [], [f5329])).
fof(f5329, plain, ($false | (~ spl72_99 | ~ spl72_104 | spl72_383)), inference(subsumption_resolution, [], [f5328, f1181])).
fof(f5328, plain, (~ (e22 = op2(e21, e23)) | (~ spl72_104 | spl72_383)), inference(forward_demodulation, [], [f2579, f1202])).
fof(f5304, plain, (spl72_30 | ~ spl72_25 | ~ spl72_231), inference(avatar_split_clause, [], [f5235, f1836, f833, f854])).
fof(f5235, plain, ((e11 = op1(e12, e10)) | (~ spl72_25 | ~ spl72_231)), inference(backward_demodulation, [], [f1837, f835])).
fof(f1837, plain, ((e11 = op1(e12, op1(e12, e11))) | ~ spl72_231), inference(avatar_component_clause, [], [f1836])).
fof(f5295, plain, (~ spl72_223 | ~ spl72_134 | spl72_212), inference(avatar_split_clause, [], [f5294, f1744, f1363, f1796])).
fof(f5294, plain, (~ (e12 = op1(op1(e12, e13), e13)) | (~ spl72_134 | spl72_212)), inference(forward_demodulation, [], [f1745, f1365])).
fof(f1745, plain, (~ (e12 = op1(op1(e13, e12), e13)) | spl72_212), inference(avatar_component_clause, [], [f1744])).
fof(f5287, plain, (~ spl72_5 | ~ spl72_7), inference(avatar_contradiction_clause, [], [f5286])).
fof(f5286, plain, ($false | (~ spl72_5 | ~ spl72_7)), inference(subsumption_resolution, [], [f5285, f361])).
fof(f361, plain, ~ (e10 = e12), inference(cnf_transformation, [], [f7])).
fof(f5285, plain, ((e10 = e12) | (~ spl72_5 | ~ spl72_7)), inference(backward_demodulation, [], [f758, f750])).
fof(f758, plain, ((e12 = op1(e13, e12)) | ~ spl72_7), inference(avatar_component_clause, [], [f756])).
fof(f5283, plain, (~ spl72_10 | spl72_217), inference(avatar_contradiction_clause, [], [f5282])).
fof(f5282, plain, ($false | (~ spl72_10 | spl72_217)), inference(subsumption_resolution, [], [f5274, f771])).
fof(f5274, plain, (~ (e11 = op1(e13, e11)) | (~ spl72_10 | spl72_217)), inference(backward_demodulation, [], [f1770, f771])).
fof(f1770, plain, (~ (e11 = op1(e13, op1(e13, e11))) | spl72_217), inference(avatar_component_clause, [], [f1768])).
fof(f1768, plain, (spl72_217 <=> (e11 = op1(e13, op1(e13, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl72_217])])).
fof(f5248, plain, (~ spl72_19 | ~ spl72_23), inference(avatar_split_clause, [], [f5242, f824, f807])).
fof(f5242, plain, (~ (e12 = op1(e12, e13)) | ~ spl72_23), inference(backward_demodulation, [], [f305, f826])).
fof(f305, plain, ~ (op1(e12, e12) = op1(e12, e13)), inference(cnf_transformation, [], [f5])).
fof(f5239, plain, (~ spl72_25 | ~ spl72_54 | spl72_230), inference(avatar_contradiction_clause, [], [f5238])).
fof(f5238, plain, ($false | (~ spl72_25 | ~ spl72_54 | spl72_230)), inference(subsumption_resolution, [], [f5234, f958])).
fof(f5234, plain, (~ (e11 = op1(e10, e12)) | (~ spl72_25 | spl72_230)), inference(backward_demodulation, [], [f1832, f835])).
fof(f1832, plain, (~ (e11 = op1(op1(e12, e11), e12)) | spl72_230), inference(avatar_component_clause, [], [f1831])).
fof(f5224, plain, (~ spl72_19 | ~ spl72_35), inference(avatar_split_clause, [], [f5220, f875, f807])).
fof(f5220, plain, (~ (e12 = op1(e12, e13)) | ~ spl72_35), inference(backward_demodulation, [], [f285, f877])).
fof(f5219, plain, (~ spl72_42 | ~ spl72_43), inference(avatar_contradiction_clause, [], [f5218])).
fof(f5218, plain, ($false | (~ spl72_42 | ~ spl72_43)), inference(subsumption_resolution, [], [f5217, f363])).
fof(f363, plain, ~ (e11 = e12), inference(cnf_transformation, [], [f7])).
fof(f5217, plain, ((e11 = e12) | (~ spl72_42 | ~ spl72_43)), inference(backward_demodulation, [], [f911, f907])).
fof(f5216, plain, (~ spl72_122 | ~ spl72_124), inference(avatar_contradiction_clause, [], [f5215])).
fof(f5215, plain, ($false | (~ spl72_122 | ~ spl72_124)), inference(subsumption_resolution, [], [f5214, f370])).
fof(f370, plain, ~ (e21 = e23), inference(cnf_transformation, [], [f8])).
fof(f5214, plain, ((e21 = e23) | (~ spl72_122 | ~ spl72_124)), inference(backward_demodulation, [], [f1287, f1279])).
fof(f1279, plain, ((e21 = op2(e20, e21)) | ~ spl72_122), inference(avatar_component_clause, [], [f1277])).
fof(f1277, plain, (spl72_122 <=> (e21 = op2(e20, e21))), introduced(avatar_definition, [new_symbols(naming, [spl72_122])])).
fof(f5207, plain, (~ spl72_93 | ~ spl72_127 | spl72_373), inference(avatar_contradiction_clause, [], [f5206])).
fof(f5206, plain, ($false | (~ spl72_93 | ~ spl72_127 | spl72_373)), inference(subsumption_resolution, [], [f5205, f1300])).
fof(f5205, plain, (~ (op2(e20, e20) = e22) | (~ spl72_93 | spl72_373)), inference(forward_demodulation, [], [f2531, f1156])).
fof(f2531, plain, (~ (e22 = op2(op2(e22, e20), e20)) | spl72_373), inference(avatar_component_clause, [], [f2529])).
fof(f2529, plain, (spl72_373 <=> (e22 = op2(op2(e22, e20), e20))), introduced(avatar_definition, [new_symbols(naming, [spl72_373])])).
fof(f5201, plain, (~ spl72_93 | spl72_375), inference(avatar_contradiction_clause, [], [f5200])).
fof(f5200, plain, ($false | (~ spl72_93 | spl72_375)), inference(subsumption_resolution, [], [f5199, f1156])).
fof(f5199, plain, (~ (e20 = op2(e22, e20)) | (~ spl72_93 | spl72_375)), inference(forward_demodulation, [], [f2541, f1156])).
fof(f5195, plain, (~ spl72_35 | ~ spl72_43), inference(avatar_split_clause, [], [f5077, f909, f875])).
fof(f5077, plain, (~ (e12 = op1(e11, e13)) | ~ spl72_43), inference(backward_demodulation, [], [f298, f911])).
fof(f5174, plain, (~ spl72_10 | ~ spl72_14), inference(avatar_split_clause, [], [f5119, f786, f769])).
fof(f5119, plain, (~ (e11 = op1(e13, e11)) | ~ spl72_14), inference(backward_demodulation, [], [f306, f788])).
fof(f306, plain, ~ (op1(e13, e10) = op1(e13, e11)), inference(cnf_transformation, [], [f5])).
fof(f5130, plain, (~ spl72_14 | spl72_33 | ~ spl72_220), inference(avatar_contradiction_clause, [], [f5129])).
fof(f5129, plain, ($false | (~ spl72_14 | spl72_33 | ~ spl72_220)), inference(subsumption_resolution, [], [f5125, f868])).
fof(f868, plain, (~ (e10 = op1(e11, e13)) | spl72_33), inference(avatar_component_clause, [], [f867])).
fof(f867, plain, (spl72_33 <=> (e10 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl72_33])])).
fof(f5125, plain, ((e10 = op1(e11, e13)) | (~ spl72_14 | ~ spl72_220)), inference(backward_demodulation, [], [f1784, f788])).
fof(f1784, plain, ((e10 = op1(op1(e13, e10), e13)) | ~ spl72_220), inference(avatar_component_clause, [], [f1782])).
fof(f5127, plain, (~ spl72_2 | ~ spl72_14), inference(avatar_split_clause, [], [f5121, f786, f735])).
fof(f735, plain, (spl72_2 <=> (e11 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl72_2])])).
fof(f5121, plain, (~ (e11 = op1(e13, e13)) | ~ spl72_14), inference(backward_demodulation, [], [f308, f788])).
fof(f308, plain, ~ (op1(e13, e10) = op1(e13, e13)), inference(cnf_transformation, [], [f5])).
fof(f5111, plain, (~ spl72_16 | ~ spl72_21 | spl72_153), inference(avatar_split_clause, [], [f5106, f1455, f816, f794])).
fof(f794, plain, (spl72_16 <=> (e13 = op1(e13, e10))), introduced(avatar_definition, [new_symbols(naming, [spl72_16])])).
fof(f5106, plain, (~ (e13 = op1(e13, e10)) | (~ spl72_21 | spl72_153)), inference(backward_demodulation, [], [f1457, f818])).
fof(f1457, plain, (~ (e13 = op1(e13, op1(e12, e12))) | spl72_153), inference(avatar_component_clause, [], [f1455])).
fof(f5110, plain, (~ spl72_5 | ~ spl72_21), inference(avatar_split_clause, [], [f5102, f816, f748])).
fof(f5102, plain, (~ (e10 = op1(e13, e12)) | ~ spl72_21), inference(backward_demodulation, [], [f281, f818])).
fof(f281, plain, ~ (op1(e12, e12) = op1(e13, e12)), inference(cnf_transformation, [], [f5])).
fof(f5101, plain, (~ spl72_16 | ~ spl72_32), inference(avatar_split_clause, [], [f5099, f862, f794])).
fof(f5099, plain, (~ (e13 = op1(e13, e10)) | ~ spl72_32), inference(backward_demodulation, [], [f269, f864])).
fof(f5094, plain, (~ spl72_2 | ~ spl72_34), inference(avatar_split_clause, [], [f5089, f871, f735])).
fof(f5089, plain, (~ (e11 = op1(e13, e13)) | ~ spl72_34), inference(backward_demodulation, [], [f286, f873])).
fof(f286, plain, ~ (op1(e11, e13) = op1(e13, e13)), inference(cnf_transformation, [], [f5])).
fof(f5088, plain, (~ spl72_36 | ~ spl72_40), inference(avatar_split_clause, [], [f5085, f896, f879])).
fof(f5085, plain, (~ (e13 = op1(e11, e13)) | ~ spl72_40), inference(backward_demodulation, [], [f299, f898])).
fof(f5080, plain, (~ spl72_11 | ~ spl72_43), inference(avatar_split_clause, [], [f5075, f909, f773])).
fof(f5075, plain, (~ (e12 = op1(e13, e11)) | ~ spl72_43), inference(backward_demodulation, [], [f274, f911])).
fof(f5050, plain, (spl72_122 | ~ spl72_110 | ~ spl72_329), inference(avatar_split_clause, [], [f5049, f2315, f1226, f1277])).
fof(f5049, plain, ((e21 = op2(e20, e21)) | (~ spl72_110 | ~ spl72_329)), inference(forward_demodulation, [], [f2317, f1228])).
fof(f5040, plain, (~ spl72_90 | spl72_371), inference(avatar_contradiction_clause, [], [f5039])).
fof(f5039, plain, ($false | (~ spl72_90 | spl72_371)), inference(subsumption_resolution, [], [f5038, f1143])).
fof(f5038, plain, (~ (e21 = op2(e22, e21)) | (~ spl72_90 | spl72_371)), inference(forward_demodulation, [], [f2522, f1143])).
fof(f5035, plain, (~ spl72_60 | spl72_526), inference(avatar_contradiction_clause, [], [f5034])).
fof(f5034, plain, ($false | (~ spl72_60 | spl72_526)), inference(trivial_inequality_removal, [], [f5033])).
fof(f5033, plain, (~ (h1(e13) = h1(e13)) | (~ spl72_60 | spl72_526)), inference(forward_demodulation, [], [f3408, f983])).
fof(f3408, plain, (~ (h1(e13) = h1(op1(e10, e11))) | spl72_526), inference(avatar_component_clause, [], [f3406])).
fof(f3406, plain, (spl72_526 <=> (h1(e13) = h1(op1(e10, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl72_526])])).
fof(f5032, plain, (~ spl72_42 | ~ spl72_26), inference(avatar_split_clause, [], [f5031, f837, f905])).
fof(f5031, plain, (~ (e11 = op1(e11, e11)) | ~ spl72_26), inference(forward_demodulation, [], [f273, f839])).
fof(f5025, plain, (~ spl72_3 | ~ spl72_19), inference(avatar_split_clause, [], [f5024, f807, f739])).
fof(f5024, plain, (~ (e12 = op1(e13, e13)) | ~ spl72_19), inference(forward_demodulation, [], [f287, f809])).
fof(f287, plain, ~ (op1(e12, e13) = op1(e13, e13)), inference(cnf_transformation, [], [f5])).
fof(f5022, plain, (~ spl72_5 | ~ spl72_54 | ~ spl72_211), inference(avatar_contradiction_clause, [], [f5021])).
fof(f5021, plain, ($false | (~ spl72_5 | ~ spl72_54 | ~ spl72_211)), inference(subsumption_resolution, [], [f5020, f364])).
fof(f5020, plain, ((e11 = e13) | (~ spl72_5 | ~ spl72_54 | ~ spl72_211)), inference(forward_demodulation, [], [f5017, f958])).
fof(f5017, plain, ((e13 = op1(e10, e12)) | (~ spl72_5 | ~ spl72_211)), inference(backward_demodulation, [], [f1740, f750])).
fof(f1740, plain, ((e13 = op1(op1(e13, e12), e12)) | ~ spl72_211), inference(avatar_component_clause, [], [f1739])).
fof(f1739, plain, (spl72_211 <=> (e13 = op1(op1(e13, e12), e12))), introduced(avatar_definition, [new_symbols(naming, [spl72_211])])).
fof(f5011, plain, (spl72_1 | ~ spl72_16 | ~ spl72_220), inference(avatar_contradiction_clause, [], [f5010])).
fof(f5010, plain, ($false | (spl72_1 | ~ spl72_16 | ~ spl72_220)), inference(subsumption_resolution, [], [f5005, f732])).
fof(f732, plain, (~ (e10 = op1(e13, e13)) | spl72_1), inference(avatar_component_clause, [], [f731])).
fof(f731, plain, (spl72_1 <=> (e10 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl72_1])])).
fof(f5005, plain, ((e10 = op1(e13, e13)) | (~ spl72_16 | ~ spl72_220)), inference(backward_demodulation, [], [f1784, f796])).
fof(f796, plain, ((e13 = op1(e13, e10)) | ~ spl72_16), inference(avatar_component_clause, [], [f794])).
fof(f5009, plain, (~ spl72_16 | spl72_219), inference(avatar_contradiction_clause, [], [f5008])).
fof(f5008, plain, ($false | (~ spl72_16 | spl72_219)), inference(subsumption_resolution, [], [f5004, f796])).
fof(f5004, plain, (~ (e13 = op1(e13, e10)) | (~ spl72_16 | spl72_219)), inference(backward_demodulation, [], [f1779, f796])).
fof(f1779, plain, (~ (e13 = op1(op1(e13, e10), e10)) | spl72_219), inference(avatar_component_clause, [], [f1777])).
fof(f1777, plain, (spl72_219 <=> (e13 = op1(op1(e13, e10), e10))), introduced(avatar_definition, [new_symbols(naming, [spl72_219])])).
fof(f5007, plain, (~ spl72_8 | ~ spl72_16), inference(avatar_split_clause, [], [f5003, f794, f760])).
fof(f5003, plain, (~ (e13 = op1(e13, e12)) | ~ spl72_16), inference(backward_demodulation, [], [f307, f796])).
fof(f307, plain, ~ (op1(e13, e10) = op1(e13, e12)), inference(cnf_transformation, [], [f5])).
fof(f4993, plain, (~ spl72_13 | ~ spl72_29), inference(avatar_split_clause, [], [f4989, f850, f782])).
fof(f4989, plain, (~ (e10 = op1(e13, e10)) | ~ spl72_29), inference(backward_demodulation, [], [f269, f852])).
fof(f4992, plain, (~ spl72_21 | ~ spl72_29), inference(avatar_split_clause, [], [f4987, f850, f816])).
fof(f4987, plain, (~ (e10 = op1(e12, e12)) | ~ spl72_29), inference(backward_demodulation, [], [f301, f852])).
fof(f4978, plain, (~ spl72_87 | spl72_294), inference(avatar_contradiction_clause, [], [f4977])).
fof(f4977, plain, ($false | (~ spl72_87 | spl72_294)), inference(subsumption_resolution, [], [f4976, f1130])).
fof(f1130, plain, ((e22 = op2(e22, e22)) | ~ spl72_87), inference(avatar_component_clause, [], [f1128])).
fof(f4976, plain, (~ (e22 = op2(e22, e22)) | (~ spl72_87 | spl72_294)), inference(forward_demodulation, [], [f2146, f1130])).
fof(f2146, plain, (~ (e22 = op2(e22, op2(e22, e22))) | spl72_294), inference(avatar_component_clause, [], [f2144])).
fof(f4967, plain, (~ spl72_84 | spl72_365), inference(avatar_contradiction_clause, [], [f4966])).
fof(f4966, plain, ($false | (~ spl72_84 | spl72_365)), inference(subsumption_resolution, [], [f4965, f1117])).
fof(f1117, plain, ((e23 = op2(e22, e23)) | ~ spl72_84), inference(avatar_component_clause, [], [f1115])).
fof(f4965, plain, (~ (e23 = op2(e22, e23)) | (~ spl72_84 | spl72_365)), inference(forward_demodulation, [], [f2492, f1117])).
fof(f2492, plain, (~ (e23 = op2(e22, op2(e22, e23))) | spl72_365), inference(avatar_component_clause, [], [f2490])).
fof(f2490, plain, (spl72_365 <=> (e23 = op2(e22, op2(e22, e23)))), introduced(avatar_definition, [new_symbols(naming, [spl72_365])])).
fof(f4949, plain, (~ spl72_14 | ~ spl72_46), inference(avatar_split_clause, [], [f4948, f922, f786])).
fof(f4948, plain, (~ (e11 = op1(e13, e10)) | ~ spl72_46), inference(forward_demodulation, [], [f268, f924])).
fof(f4933, plain, (~ spl72_94 | ~ spl72_110), inference(avatar_split_clause, [], [f4930, f1226, f1158])).
fof(f4930, plain, (~ (e21 = op2(e22, e20)) | ~ spl72_110), inference(backward_demodulation, [], [f315, f1228])).
fof(f4923, plain, (~ spl72_13 | ~ spl72_60 | ~ spl72_260), inference(avatar_contradiction_clause, [], [f4922])).
fof(f4922, plain, ($false | (~ spl72_13 | ~ spl72_60 | ~ spl72_260)), inference(subsumption_resolution, [], [f4921, f360])).
fof(f4921, plain, ((e10 = e11) | (~ spl72_13 | ~ spl72_60 | ~ spl72_260)), inference(forward_demodulation, [], [f4920, f784])).
fof(f4920, plain, ((e11 = op1(e13, e10)) | (~ spl72_60 | ~ spl72_260)), inference(forward_demodulation, [], [f1977, f983])).
fof(f4902, plain, (~ spl72_80 | spl72_359), inference(avatar_contradiction_clause, [], [f4901])).
fof(f4901, plain, ($false | (~ spl72_80 | spl72_359)), inference(subsumption_resolution, [], [f4900, f1100])).
fof(f1100, plain, ((e23 = op2(e23, e20)) | ~ spl72_80), inference(avatar_component_clause, [], [f1098])).
fof(f4900, plain, (~ (e23 = op2(e23, e20)) | (~ spl72_80 | spl72_359)), inference(forward_demodulation, [], [f2463, f1100])).
fof(f4899, plain, (spl72_65 | ~ spl72_80 | ~ spl72_360), inference(avatar_contradiction_clause, [], [f4898])).
fof(f4898, plain, ($false | (spl72_65 | ~ spl72_80 | ~ spl72_360)), inference(subsumption_resolution, [], [f4897, f1036])).
fof(f1036, plain, (~ (e20 = op2(e23, e23)) | spl72_65), inference(avatar_component_clause, [], [f1035])).
fof(f1035, plain, (spl72_65 <=> (e20 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl72_65])])).
fof(f4897, plain, ((e20 = op2(e23, e23)) | (~ spl72_80 | ~ spl72_360)), inference(forward_demodulation, [], [f2468, f1100])).
fof(f2468, plain, ((e20 = op2(op2(e23, e20), e23)) | ~ spl72_360), inference(avatar_component_clause, [], [f2466])).
fof(f2466, plain, (spl72_360 <=> (e20 = op2(op2(e23, e20), e23))), introduced(avatar_definition, [new_symbols(naming, [spl72_360])])).
fof(f4839, plain, (~ spl72_26 | spl72_231), inference(avatar_contradiction_clause, [], [f4838])).
fof(f4838, plain, ($false | (~ spl72_26 | spl72_231)), inference(subsumption_resolution, [], [f4834, f839])).
fof(f4834, plain, (~ (e11 = op1(e12, e11)) | (~ spl72_26 | spl72_231)), inference(backward_demodulation, [], [f1838, f839])).
fof(f4774, plain, (spl72_79 | ~ spl72_69 | ~ spl72_353), inference(avatar_split_clause, [], [f4643, f2433, f1052, f1094])).
fof(f1094, plain, (spl72_79 <=> (e22 = op2(e23, e20))), introduced(avatar_definition, [new_symbols(naming, [spl72_79])])).
fof(f2433, plain, (spl72_353 <=> (e22 = op2(e23, op2(e23, e22)))), introduced(avatar_definition, [new_symbols(naming, [spl72_353])])).
fof(f4643, plain, ((e22 = op2(e23, e20)) | (~ spl72_69 | ~ spl72_353)), inference(forward_demodulation, [], [f2434, f1054])).
fof(f2434, plain, ((e22 = op2(e23, op2(e23, e22))) | ~ spl72_353), inference(avatar_component_clause, [], [f2433])).
fof(f4772, plain, (~ spl72_78 | ~ spl72_415), inference(avatar_split_clause, [], [f4771, f2784, f1090])).
fof(f4771, plain, (~ (e21 = op2(e23, e20)) | ~ spl72_415), inference(forward_demodulation, [], [f2748, f2785])).
fof(f4768, plain, (~ spl72_42 | ~ spl72_46), inference(avatar_split_clause, [], [f4767, f922, f905])).
fof(f4767, plain, (~ (e11 = op1(e11, e11)) | ~ spl72_46), inference(forward_demodulation, [], [f294, f924])).
fof(f4747, plain, (~ spl72_19 | spl72_223), inference(avatar_contradiction_clause, [], [f4746])).
fof(f4746, plain, ($false | (~ spl72_19 | spl72_223)), inference(subsumption_resolution, [], [f4742, f809])).
fof(f4742, plain, (~ (e12 = op1(e12, e13)) | (~ spl72_19 | spl72_223)), inference(backward_demodulation, [], [f1798, f809])).
fof(f4736, plain, (~ spl72_20 | ~ spl72_24), inference(avatar_split_clause, [], [f4731, f828, f811])).
fof(f4731, plain, (~ (e13 = op1(e12, e13)) | ~ spl72_24), inference(backward_demodulation, [], [f305, f830])).
fof(f4728, plain, (~ spl72_37 | ~ spl72_39), inference(avatar_contradiction_clause, [], [f4727])).
fof(f4727, plain, ($false | (~ spl72_37 | ~ spl72_39)), inference(subsumption_resolution, [], [f4726, f361])).
fof(f4726, plain, ((e10 = e12) | (~ spl72_37 | ~ spl72_39)), inference(forward_demodulation, [], [f894, f886])).
fof(f4721, plain, (~ spl72_63 | ~ spl72_64), inference(avatar_contradiction_clause, [], [f4720])).
fof(f4720, plain, ($false | (~ spl72_63 | ~ spl72_64)), inference(subsumption_resolution, [], [f4719, f365])).
fof(f4719, plain, ((e12 = e13) | (~ spl72_63 | ~ spl72_64)), inference(forward_demodulation, [], [f1000, f996])).
fof(f1000, plain, ((op1(e10, e10) = e13) | ~ spl72_64), inference(avatar_component_clause, [], [f998])).
fof(f998, plain, (spl72_64 <=> (op1(e10, e10) = e13)), introduced(avatar_definition, [new_symbols(naming, [spl72_64])])).
fof(f4713, plain, (~ spl72_82 | ~ spl72_94), inference(avatar_split_clause, [], [f4711, f1158, f1107])).
fof(f4711, plain, (~ (e21 = op2(e22, e23)) | ~ spl72_94), inference(backward_demodulation, [], [f350, f1160])).
fof(f4709, plain, (~ spl72_127 | ~ spl72_128), inference(avatar_contradiction_clause, [], [f4708])).
fof(f4708, plain, ($false | (~ spl72_127 | ~ spl72_128)), inference(subsumption_resolution, [], [f4707, f371])).
fof(f4707, plain, ((e22 = e23) | (~ spl72_127 | ~ spl72_128)), inference(forward_demodulation, [], [f1304, f1300])).
fof(f1304, plain, ((op2(e20, e20) = e23) | ~ spl72_128), inference(avatar_component_clause, [], [f1302])).
fof(f1302, plain, (spl72_128 <=> (op2(e20, e20) = e23)), introduced(avatar_definition, [new_symbols(naming, [spl72_128])])).
fof(f4698, plain, (~ spl72_75 | ~ spl72_87 | spl72_277), inference(avatar_contradiction_clause, [], [f4697])).
fof(f4697, plain, ($false | (~ spl72_75 | ~ spl72_87 | spl72_277)), inference(subsumption_resolution, [], [f4696, f1130])).
fof(f4696, plain, (~ (e22 = op2(e22, e22)) | (~ spl72_75 | spl72_277)), inference(forward_demodulation, [], [f2063, f1079])).
fof(f2063, plain, (~ (e22 = op2(e22, op2(e23, e21))) | spl72_277), inference(avatar_component_clause, [], [f2061])).
fof(f2061, plain, (spl72_277 <=> (e22 = op2(e22, op2(e23, e21)))), introduced(avatar_definition, [new_symbols(naming, [spl72_277])])).
fof(f4650, plain, (~ spl72_69 | ~ spl72_118 | ~ spl72_351), inference(avatar_contradiction_clause, [], [f4649])).
fof(f4649, plain, ($false | (~ spl72_69 | ~ spl72_118 | ~ spl72_351)), inference(subsumption_resolution, [], [f4648, f370])).
fof(f4648, plain, ((e21 = e23) | (~ spl72_69 | ~ spl72_118 | ~ spl72_351)), inference(forward_demodulation, [], [f4647, f1262])).
fof(f4647, plain, ((e23 = op2(e20, e22)) | (~ spl72_69 | ~ spl72_351)), inference(forward_demodulation, [], [f2424, f1054])).
fof(f2424, plain, ((e23 = op2(op2(e23, e22), e22)) | ~ spl72_351), inference(avatar_component_clause, [], [f2423])).
fof(f4637, plain, (spl72_82 | ~ spl72_75 | ~ spl72_356), inference(avatar_split_clause, [], [f4574, f2447, f1077, f1107])).
fof(f2447, plain, (spl72_356 <=> (e21 = op2(op2(e23, e21), e23))), introduced(avatar_definition, [new_symbols(naming, [spl72_356])])).
fof(f4574, plain, ((e21 = op2(e22, e23)) | (~ spl72_75 | ~ spl72_356)), inference(backward_demodulation, [], [f2449, f1079])).
fof(f2449, plain, ((e21 = op2(op2(e23, e21), e23)) | ~ spl72_356), inference(avatar_component_clause, [], [f2447])).
fof(f4629, plain, (~ spl72_18 | ~ spl72_20), inference(avatar_contradiction_clause, [], [f4628])).
fof(f4628, plain, ($false | (~ spl72_18 | ~ spl72_20)), inference(subsumption_resolution, [], [f4627, f364])).
fof(f4627, plain, ((e11 = e13) | (~ spl72_18 | ~ spl72_20)), inference(backward_demodulation, [], [f813, f805])).
fof(f4614, plain, (~ spl72_26 | ~ spl72_30), inference(avatar_split_clause, [], [f4610, f854, f837])).
fof(f4610, plain, (~ (e11 = op1(e12, e11)) | ~ spl72_30), inference(backward_demodulation, [], [f300, f856])).
fof(f4581, plain, (spl72_472 | ~ spl72_68 | ~ spl72_471), inference(avatar_split_clause, [], [f4580, f3076, f1047, f3091])).
fof(f4580, plain, ((e23 = h4(e11)) | (~ spl72_68 | ~ spl72_471)), inference(backward_demodulation, [], [f4492, f1049])).
fof(f4492, plain, ((op2(e23, e23) = h4(e11)) | ~ spl72_471), inference(backward_demodulation, [], [f2752, f3077])).
fof(f4566, plain, (spl72_75 | ~ spl72_82 | ~ spl72_274 | ~ spl72_353), inference(avatar_split_clause, [], [f4561, f2433, f2047, f1107, f1077])).
fof(f2047, plain, (spl72_274 <=> (op2(e22, e23) = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl72_274])])).
fof(f4561, plain, ((e22 = op2(e23, e21)) | (~ spl72_82 | ~ spl72_274 | ~ spl72_353)), inference(backward_demodulation, [], [f4480, f1109])).
fof(f4480, plain, ((e22 = op2(e23, op2(e22, e23))) | (~ spl72_274 | ~ spl72_353)), inference(forward_demodulation, [], [f2434, f2049])).
fof(f2049, plain, ((op2(e22, e23) = op2(e23, e22)) | ~ spl72_274), inference(avatar_component_clause, [], [f2047])).
fof(f4564, plain, (spl72_70 | ~ spl72_82 | ~ spl72_274), inference(avatar_contradiction_clause, [], [f4563])).
fof(f4563, plain, ($false | (spl72_70 | ~ spl72_82 | ~ spl72_274)), inference(subsumption_resolution, [], [f4559, f1057])).
fof(f4559, plain, ((e21 = op2(e23, e22)) | (~ spl72_82 | ~ spl72_274)), inference(backward_demodulation, [], [f2049, f1109])).
fof(f4551, plain, (~ spl72_72 | ~ spl72_104), inference(avatar_split_clause, [], [f4548, f1200, f1064])).
fof(f4548, plain, (~ (e23 = op2(e23, e22)) | ~ spl72_104), inference(backward_demodulation, [], [f328, f1202])).
fof(f4542, plain, (~ spl72_101 | ~ spl72_109), inference(avatar_split_clause, [], [f4536, f1222, f1188])).
fof(f4536, plain, (~ (e20 = op2(e21, e22)) | ~ spl72_109), inference(backward_demodulation, [], [f343, f1224])).
fof(f4530, plain, (~ spl72_23 | ~ spl72_11 | spl72_137), inference(avatar_split_clause, [], [f4529, f1377, f773, f824])).
fof(f1377, plain, (spl72_137 <=> (e12 = op1(e12, op1(e13, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl72_137])])).
fof(f4529, plain, (~ (e12 = op1(e12, e12)) | (~ spl72_11 | spl72_137)), inference(forward_demodulation, [], [f1379, f775])).
fof(f1379, plain, (~ (e12 = op1(e12, op1(e13, e11))) | spl72_137), inference(avatar_component_clause, [], [f1377])).
fof(f4524, plain, (spl72_18 | ~ spl72_11 | ~ spl72_216), inference(avatar_split_clause, [], [f4523, f1763, f773, f803])).
fof(f4523, plain, ((e11 = op1(e12, e13)) | (~ spl72_11 | ~ spl72_216)), inference(forward_demodulation, [], [f1765, f775])).
fof(f4503, plain, (~ spl72_68 | spl72_404 | ~ spl72_471), inference(avatar_split_clause, [], [f4495, f3076, f2680, f1047])).
fof(f2680, plain, (spl72_404 <=> (e23 = op2(e23, op2(e23, e23)))), introduced(avatar_definition, [new_symbols(naming, [spl72_404])])).
fof(f4495, plain, (~ (e23 = op2(e23, e23)) | (spl72_404 | ~ spl72_471)), inference(backward_demodulation, [], [f3442, f3077])).
fof(f3442, plain, (~ (e23 = op2(e23, h4(e12))) | spl72_404), inference(forward_demodulation, [], [f2682, f664])).
fof(f2682, plain, (~ (e23 = op2(e23, op2(e23, e23))) | spl72_404), inference(avatar_component_clause, [], [f2680])).
fof(f4501, plain, (~ spl72_113 | spl72_407 | ~ spl72_471), inference(avatar_contradiction_clause, [], [f4500])).
fof(f4500, plain, ($false | (~ spl72_113 | spl72_407 | ~ spl72_471)), inference(subsumption_resolution, [], [f4493, f1241])).
fof(f4493, plain, (~ (e20 = op2(e20, e23)) | (spl72_407 | ~ spl72_471)), inference(backward_demodulation, [], [f3439, f3077])).
fof(f3439, plain, (~ (e20 = op2(e20, h4(e12))) | spl72_407), inference(forward_demodulation, [], [f2697, f664])).
fof(f2697, plain, (~ (e20 = op2(e20, op2(e23, e23))) | spl72_407), inference(avatar_component_clause, [], [f2695])).
fof(f4449, plain, (~ spl72_18 | ~ spl72_2), inference(avatar_split_clause, [], [f4448, f735, f803])).
fof(f4448, plain, (~ (e11 = op1(e12, e13)) | ~ spl72_2), inference(forward_demodulation, [], [f287, f737])).
fof(f737, plain, ((e11 = op1(e13, e13)) | ~ spl72_2), inference(avatar_component_clause, [], [f735])).
fof(f4411, plain, (~ spl72_13 | spl72_221), inference(avatar_contradiction_clause, [], [f4410])).
fof(f4410, plain, ($false | (~ spl72_13 | spl72_221)), inference(subsumption_resolution, [], [f4409, f784])).
fof(f4409, plain, (~ (e10 = op1(e13, e10)) | (~ spl72_13 | spl72_221)), inference(forward_demodulation, [], [f1789, f784])).
fof(f4393, plain, (~ spl72_74 | spl72_357), inference(avatar_contradiction_clause, [], [f4392])).
fof(f4392, plain, ($false | (~ spl72_74 | spl72_357)), inference(subsumption_resolution, [], [f4391, f1075])).
fof(f4391, plain, (~ (e21 = op2(e23, e21)) | (~ spl72_74 | spl72_357)), inference(forward_demodulation, [], [f2454, f1075])).
fof(f2454, plain, (~ (e21 = op2(e23, op2(e23, e21))) | spl72_357), inference(avatar_component_clause, [], [f2452])).
fof(f2452, plain, (spl72_357 <=> (e21 = op2(e23, op2(e23, e21)))), introduced(avatar_definition, [new_symbols(naming, [spl72_357])])).
fof(f4382, plain, (~ spl72_110 | ~ spl72_98), inference(avatar_split_clause, [], [f4381, f1175, f1226])).
fof(f4381, plain, (~ (e21 = op2(e21, e20)) | ~ spl72_98), inference(forward_demodulation, [], [f344, f1177])).
fof(f4380, plain, (~ spl72_72 | spl72_404 | ~ spl72_411), inference(avatar_split_clause, [], [f4271, f2764, f2680, f1064])).
fof(f4271, plain, (~ (e23 = op2(e23, e22)) | (spl72_404 | ~ spl72_411)), inference(backward_demodulation, [], [f3442, f2765])).
fof(f4358, plain, (~ spl72_4 | ~ spl72_8), inference(avatar_split_clause, [], [f4354, f760, f743])).
fof(f4354, plain, (~ (e13 = op1(e13, e13)) | ~ spl72_8), inference(backward_demodulation, [], [f311, f762])).
fof(f311, plain, ~ (op1(e13, e12) = op1(e13, e13)), inference(cnf_transformation, [], [f5])).
fof(f4349, plain, (~ spl72_3 | ~ spl72_11), inference(avatar_split_clause, [], [f4344, f773, f739])).
fof(f4344, plain, (~ (e12 = op1(e13, e13)) | ~ spl72_11), inference(backward_demodulation, [], [f310, f775])).
fof(f310, plain, ~ (op1(e13, e11) = op1(e13, e13)), inference(cnf_transformation, [], [f5])).
fof(f4341, plain, (~ spl72_9 | ~ spl72_13), inference(avatar_split_clause, [], [f4339, f782, f765])).
fof(f4339, plain, (~ (e10 = op1(e13, e11)) | ~ spl72_13), inference(backward_demodulation, [], [f306, f784])).
fof(f4331, plain, (~ spl72_39 | ~ spl72_40), inference(avatar_contradiction_clause, [], [f4330])).
fof(f4330, plain, ($false | (~ spl72_39 | ~ spl72_40)), inference(subsumption_resolution, [], [f4329, f365])).
fof(f4329, plain, ((e12 = e13) | (~ spl72_39 | ~ spl72_40)), inference(backward_demodulation, [], [f898, f894])).
fof(f4328, plain, (~ spl72_9 | ~ spl72_41), inference(avatar_split_clause, [], [f4325, f901, f765])).
fof(f4325, plain, (~ (e10 = op1(e13, e11)) | ~ spl72_41), inference(backward_demodulation, [], [f274, f903])).
fof(f4284, plain, (~ spl72_121 | ~ spl72_124), inference(avatar_contradiction_clause, [], [f4283])).
fof(f4283, plain, ($false | (~ spl72_121 | ~ spl72_124)), inference(subsumption_resolution, [], [f4282, f368])).
fof(f368, plain, ~ (e20 = e23), inference(cnf_transformation, [], [f8])).
fof(f4282, plain, ((e20 = e23) | (~ spl72_121 | ~ spl72_124)), inference(backward_demodulation, [], [f1287, f1275])).
fof(f1275, plain, ((e20 = op2(e20, e21)) | ~ spl72_121), inference(avatar_component_clause, [], [f1273])).
fof(f1273, plain, (spl72_121 <=> (e20 = op2(e20, e21))), introduced(avatar_definition, [new_symbols(naming, [spl72_121])])).
fof(f4281, plain, (~ spl72_87 | spl72_405 | ~ spl72_411), inference(avatar_contradiction_clause, [], [f4280])).
fof(f4280, plain, ($false | (~ spl72_87 | spl72_405 | ~ spl72_411)), inference(subsumption_resolution, [], [f4270, f1130])).
fof(f4270, plain, (~ (e22 = op2(e22, e22)) | (spl72_405 | ~ spl72_411)), inference(backward_demodulation, [], [f3441, f2765])).
fof(f3441, plain, (~ (e22 = op2(e22, h4(e12))) | spl72_405), inference(forward_demodulation, [], [f2687, f664])).
fof(f4250, plain, (spl72_423 | ~ spl72_87), inference(avatar_split_clause, [], [f4147, f1128, f2824])).
fof(f4147, plain, ((e22 = h3(e12)) | ~ spl72_87), inference(backward_demodulation, [], [f660, f1130])).
fof(f4249, plain, (spl72_411 | ~ spl72_67), inference(avatar_split_clause, [], [f4179, f1043, f2764])).
fof(f4179, plain, ((e22 = h4(e12)) | ~ spl72_67), inference(backward_demodulation, [], [f664, f1045])).
fof(f4245, plain, (~ spl72_116 | ~ spl72_515), inference(avatar_split_clause, [], [f4118, f3362, f1251])).
fof(f1251, plain, (spl72_116 <=> (e23 = op2(e20, e23))), introduced(avatar_definition, [new_symbols(naming, [spl72_116])])).
fof(f4118, plain, (~ (e23 = op2(e20, e23)) | ~ spl72_515), inference(forward_demodulation, [], [f2718, f3363])).
fof(f2718, plain, ~ (op2(e20, e23) = h1(e13)), inference(backward_demodulation, [], [f340, f2712])).
fof(f340, plain, ~ (op2(e20, e21) = op2(e20, e23)), inference(cnf_transformation, [], [f6])).
fof(f4244, plain, (~ spl72_115 | ~ spl72_447), inference(avatar_split_clause, [], [f4120, f2946, f1247])).
fof(f1247, plain, (spl72_115 <=> (e22 = op2(e20, e23))), introduced(avatar_definition, [new_symbols(naming, [spl72_115])])).
fof(f4120, plain, (~ (e22 = op2(e20, e23)) | ~ spl72_447), inference(forward_demodulation, [], [f2723, f2947])).
fof(f2723, plain, ~ (op2(e20, e23) = h1(e12)), inference(backward_demodulation, [], [f338, f652])).
fof(f338, plain, ~ (op2(e20, e20) = op2(e20, e23)), inference(cnf_transformation, [], [f6])).
fof(f4206, plain, (~ spl72_10 | ~ spl72_26), inference(avatar_split_clause, [], [f4205, f837, f769])).
fof(f4205, plain, (~ (e11 = op1(e13, e11)) | ~ spl72_26), inference(forward_demodulation, [], [f275, f839])).
fof(f4203, plain, (~ spl72_1 | ~ spl72_49), inference(avatar_split_clause, [], [f4202, f935, f731])).
fof(f4202, plain, (~ (e10 = op1(e13, e13)) | ~ spl72_49), inference(forward_demodulation, [], [f284, f937])).
fof(f284, plain, ~ (op1(e10, e13) = op1(e13, e13)), inference(cnf_transformation, [], [f5])).
fof(f4199, plain, (~ spl72_3 | ~ spl72_4), inference(avatar_contradiction_clause, [], [f4198])).
fof(f4198, plain, ($false | (~ spl72_3 | ~ spl72_4)), inference(subsumption_resolution, [], [f4197, f365])).
fof(f4197, plain, ((e12 = e13) | (~ spl72_3 | ~ spl72_4)), inference(backward_demodulation, [], [f745, f741])).
fof(f4185, plain, (~ spl72_21 | ~ spl72_23), inference(avatar_contradiction_clause, [], [f4184])).
fof(f4184, plain, ($false | (~ spl72_21 | ~ spl72_23)), inference(subsumption_resolution, [], [f4183, f361])).
fof(f4183, plain, ((e10 = e12) | (~ spl72_21 | ~ spl72_23)), inference(forward_demodulation, [], [f826, f818])).
fof(f4177, plain, (~ spl72_71 | spl72_353), inference(avatar_contradiction_clause, [], [f4176])).
fof(f4176, plain, ($false | (~ spl72_71 | spl72_353)), inference(subsumption_resolution, [], [f4169, f1062])).
fof(f4169, plain, (~ (e22 = op2(e23, e22)) | (~ spl72_71 | spl72_353)), inference(backward_demodulation, [], [f2435, f1062])).
fof(f2435, plain, (~ (e22 = op2(e23, op2(e23, e22))) | spl72_353), inference(avatar_component_clause, [], [f2433])).
fof(f4155, plain, (~ spl72_82 | ~ spl72_84), inference(avatar_contradiction_clause, [], [f4154])).
fof(f4154, plain, ($false | (~ spl72_82 | ~ spl72_84)), inference(subsumption_resolution, [], [f4152, f370])).
fof(f4152, plain, ((e21 = e23) | (~ spl72_82 | ~ spl72_84)), inference(backward_demodulation, [], [f1117, f1109])).
fof(f4145, plain, (~ spl72_89 | ~ spl72_91), inference(avatar_contradiction_clause, [], [f4144])).
fof(f4144, plain, ($false | (~ spl72_89 | ~ spl72_91)), inference(subsumption_resolution, [], [f4143, f367])).
fof(f367, plain, ~ (e20 = e22), inference(cnf_transformation, [], [f8])).
fof(f4143, plain, ((e20 = e22) | (~ spl72_89 | ~ spl72_91)), inference(backward_demodulation, [], [f1147, f1139])).
fof(f4126, plain, (spl72_435 | ~ spl72_107), inference(avatar_split_clause, [], [f3966, f1213, f2885])).
fof(f3966, plain, ((e22 = h2(e12)) | ~ spl72_107), inference(backward_demodulation, [], [f656, f1215])).
fof(f4123, plain, (~ spl72_419 | ~ spl72_113), inference(avatar_split_clause, [], [f4122, f1239, f2804])).
fof(f2804, plain, (spl72_419 <=> (e20 = h4(e12))), introduced(avatar_definition, [new_symbols(naming, [spl72_419])])).
fof(f4122, plain, (~ (e20 = h4(e12)) | ~ spl72_113), inference(forward_demodulation, [], [f2745, f1241])).
fof(f2745, plain, ~ (op2(e20, e23) = h4(e12)), inference(backward_demodulation, [], [f332, f664])).
fof(f332, plain, ~ (op2(e20, e23) = op2(e23, e23)), inference(cnf_transformation, [], [f6])).
fof(f4115, plain, (~ spl72_81 | ~ spl72_113), inference(avatar_split_clause, [], [f4114, f1239, f1103])).
fof(f1103, plain, (spl72_81 <=> (e20 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl72_81])])).
fof(f4114, plain, (~ (e20 = op2(e22, e23)) | ~ spl72_113), inference(forward_demodulation, [], [f331, f1241])).
fof(f331, plain, ~ (op2(e20, e23) = op2(e22, e23)), inference(cnf_transformation, [], [f6])).
fof(f4112, plain, (~ spl72_111 | ~ spl72_447), inference(avatar_split_clause, [], [f4111, f2946, f1230])).
fof(f4111, plain, (~ (e22 = op2(e21, e20)) | ~ spl72_447), inference(forward_demodulation, [], [f2719, f2947])).
fof(f2719, plain, ~ (op2(e21, e20) = h1(e12)), inference(backward_demodulation, [], [f312, f652])).
fof(f312, plain, ~ (op2(e20, e20) = op2(e21, e20)), inference(cnf_transformation, [], [f6])).
fof(f4095, plain, (~ spl72_102 | ~ spl72_118), inference(avatar_contradiction_clause, [], [f4094])).
fof(f4094, plain, ($false | (~ spl72_102 | ~ spl72_118)), inference(subsumption_resolution, [], [f4093, f1262])).
fof(f4093, plain, (~ (e21 = op2(e20, e22)) | ~ spl72_102), inference(forward_demodulation, [], [f324, f1194])).
fof(f1194, plain, ((e21 = op2(e21, e22)) | ~ spl72_102), inference(avatar_component_clause, [], [f1192])).
fof(f1192, plain, (spl72_102 <=> (e21 = op2(e21, e22))), introduced(avatar_definition, [new_symbols(naming, [spl72_102])])).
fof(f324, plain, ~ (op2(e20, e22) = op2(e21, e22)), inference(cnf_transformation, [], [f6])).
fof(f4068, plain, (~ spl72_6 | ~ spl72_54), inference(avatar_split_clause, [], [f3894, f956, f752])).
fof(f3894, plain, (~ (e11 = op1(e13, e12)) | ~ spl72_54), inference(forward_demodulation, [], [f278, f958])).
fof(f278, plain, ~ (op1(e10, e12) = op1(e13, e12)), inference(cnf_transformation, [], [f5])).
fof(f4066, plain, (~ spl72_4 | spl72_264), inference(avatar_contradiction_clause, [], [f4065])).
fof(f4065, plain, ($false | (~ spl72_4 | spl72_264)), inference(subsumption_resolution, [], [f4064, f745])).
fof(f4064, plain, (~ (e13 = op1(e13, e13)) | (~ spl72_4 | spl72_264)), inference(forward_demodulation, [], [f1998, f745])).
fof(f4057, plain, (~ spl72_4 | ~ spl72_49 | spl72_267), inference(avatar_contradiction_clause, [], [f4056])).
fof(f4056, plain, ($false | (~ spl72_4 | ~ spl72_49 | spl72_267)), inference(subsumption_resolution, [], [f4055, f937])).
fof(f4055, plain, (~ (e10 = op1(e10, e13)) | (~ spl72_4 | spl72_267)), inference(forward_demodulation, [], [f2013, f745])).
fof(f2013, plain, (~ (e10 = op1(e10, op1(e13, e13))) | spl72_267), inference(avatar_component_clause, [], [f2011])).
fof(f2011, plain, (spl72_267 <=> (e10 = op1(e10, op1(e13, e13)))), introduced(avatar_definition, [new_symbols(naming, [spl72_267])])).
fof(f4042, plain, (~ spl72_2 | ~ spl72_4), inference(avatar_contradiction_clause, [], [f4041])).
fof(f4041, plain, ($false | (~ spl72_2 | ~ spl72_4)), inference(subsumption_resolution, [], [f4040, f364])).
fof(f4040, plain, ((e11 = e13) | (~ spl72_2 | ~ spl72_4)), inference(forward_demodulation, [], [f745, f737])).
fof(f4039, plain, (~ spl72_7 | spl72_213), inference(avatar_contradiction_clause, [], [f4038])).
fof(f4038, plain, ($false | (~ spl72_7 | spl72_213)), inference(subsumption_resolution, [], [f4036, f758])).
fof(f4036, plain, (~ (e12 = op1(e13, e12)) | (~ spl72_7 | spl72_213)), inference(backward_demodulation, [], [f1751, f758])).
fof(f1751, plain, (~ (e12 = op1(e13, op1(e13, e12))) | spl72_213), inference(avatar_component_clause, [], [f1749])).
fof(f1749, plain, (spl72_213 <=> (e12 = op1(e13, op1(e13, e12)))), introduced(avatar_definition, [new_symbols(naming, [spl72_213])])).
fof(f4028, plain, (~ spl72_21 | ~ spl72_24), inference(avatar_contradiction_clause, [], [f4027])).
fof(f4027, plain, ($false | (~ spl72_21 | ~ spl72_24)), inference(subsumption_resolution, [], [f4026, f362])).
fof(f4026, plain, ((e10 = e13) | (~ spl72_21 | ~ spl72_24)), inference(backward_demodulation, [], [f830, f818])).
fof(f4014, plain, (~ spl72_68 | spl72_471), inference(avatar_contradiction_clause, [], [f4013])).
fof(f4013, plain, ($false | (~ spl72_68 | spl72_471)), inference(subsumption_resolution, [], [f4012, f3078])).
fof(f3078, plain, (~ (e23 = h4(e12)) | spl72_471), inference(avatar_component_clause, [], [f3076])).
fof(f4012, plain, ((e23 = h4(e12)) | ~ spl72_68), inference(backward_demodulation, [], [f664, f1049])).
fof(f4008, plain, (~ spl72_69 | ~ spl72_113 | ~ spl72_352), inference(avatar_contradiction_clause, [], [f4007])).
fof(f4007, plain, ($false | (~ spl72_69 | ~ spl72_113 | ~ spl72_352)), inference(subsumption_resolution, [], [f4006, f367])).
fof(f4006, plain, ((e20 = e22) | (~ spl72_69 | ~ spl72_113 | ~ spl72_352)), inference(forward_demodulation, [], [f4001, f1241])).
fof(f4001, plain, ((e22 = op2(e20, e23)) | (~ spl72_69 | ~ spl72_352)), inference(backward_demodulation, [], [f2430, f1054])).
fof(f2430, plain, ((e22 = op2(op2(e23, e22), e23)) | ~ spl72_352), inference(avatar_component_clause, [], [f2428])).
fof(f2428, plain, (spl72_352 <=> (e22 = op2(op2(e23, e22), e23))), introduced(avatar_definition, [new_symbols(naming, [spl72_352])])).
fof(f3988, plain, (~ spl72_86 | spl72_427), inference(avatar_contradiction_clause, [], [f3987])).
fof(f3987, plain, ($false | (~ spl72_86 | spl72_427)), inference(subsumption_resolution, [], [f3986, f2847])).
fof(f2847, plain, (~ (e21 = h3(e12)) | spl72_427), inference(avatar_component_clause, [], [f2845])).
fof(f3986, plain, ((e21 = h3(e12)) | ~ spl72_86), inference(backward_demodulation, [], [f660, f1126])).
fof(f1126, plain, ((e21 = op2(e22, e22)) | ~ spl72_86), inference(avatar_component_clause, [], [f1124])).
fof(f3980, plain, (~ spl72_94 | ~ spl72_96), inference(avatar_contradiction_clause, [], [f3979])).
fof(f3979, plain, ($false | (~ spl72_94 | ~ spl72_96)), inference(subsumption_resolution, [], [f3978, f370])).
fof(f3978, plain, ((e21 = e23) | (~ spl72_94 | ~ spl72_96)), inference(forward_demodulation, [], [f1168, f1160])).
fof(f3971, plain, (~ spl72_103 | ~ spl72_104), inference(avatar_contradiction_clause, [], [f3970])).
fof(f3970, plain, ($false | (~ spl72_103 | ~ spl72_104)), inference(subsumption_resolution, [], [f3969, f371])).
fof(f3969, plain, ((e22 = e23) | (~ spl72_103 | ~ spl72_104)), inference(forward_demodulation, [], [f1202, f1198])).
fof(f3961, plain, (~ spl72_70 | ~ spl72_118), inference(avatar_split_clause, [], [f3960, f1260, f1056])).
fof(f3960, plain, (~ (e21 = op2(e23, e22)) | ~ spl72_118), inference(backward_demodulation, [], [f326, f1262])).
fof(f326, plain, ~ (op2(e20, e22) = op2(e23, e22)), inference(cnf_transformation, [], [f6])).
fof(f3958, plain, (~ spl72_119 | ~ spl72_447), inference(avatar_split_clause, [], [f3957, f2946, f1264])).
fof(f1264, plain, (spl72_119 <=> (e22 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl72_119])])).
fof(f3957, plain, (~ (e22 = op2(e20, e22)) | ~ spl72_447), inference(forward_demodulation, [], [f2722, f2947])).
fof(f2722, plain, ~ (op2(e20, e22) = h1(e12)), inference(backward_demodulation, [], [f337, f652])).
fof(f337, plain, ~ (op2(e20, e20) = op2(e20, e22)), inference(cnf_transformation, [], [f6])).
fof(f3956, plain, (~ spl72_120 | ~ spl72_515), inference(avatar_split_clause, [], [f3955, f3362, f1268])).
fof(f1268, plain, (spl72_120 <=> (e23 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl72_120])])).
fof(f3955, plain, (~ (e23 = op2(e20, e22)) | ~ spl72_515), inference(forward_demodulation, [], [f2717, f3363])).
fof(f2717, plain, ~ (op2(e20, e22) = h1(e13)), inference(backward_demodulation, [], [f339, f2712])).
fof(f339, plain, ~ (op2(e20, e21) = op2(e20, e22)), inference(cnf_transformation, [], [f6])).
fof(f3952, plain, (~ spl72_118 | ~ spl72_114), inference(avatar_split_clause, [], [f3951, f1243, f1260])).
fof(f1243, plain, (spl72_114 <=> (e21 = op2(e20, e23))), introduced(avatar_definition, [new_symbols(naming, [spl72_114])])).
fof(f3951, plain, (~ (e21 = op2(e20, e22)) | ~ spl72_114), inference(forward_demodulation, [], [f341, f1245])).
fof(f1245, plain, ((e21 = op2(e20, e23)) | ~ spl72_114), inference(avatar_component_clause, [], [f1243])).
fof(f341, plain, ~ (op2(e20, e22) = op2(e20, e23)), inference(cnf_transformation, [], [f6])).
fof(f3950, plain, (spl72_118 | ~ spl72_447), inference(avatar_split_clause, [], [f3812, f2946, f1260])).
fof(f3812, plain, ((e21 = op2(e20, e22)) | ~ spl72_447), inference(backward_demodulation, [], [f2724, f2947])).
fof(f3940, plain, (~ spl72_79 | ~ spl72_447), inference(avatar_split_clause, [], [f3939, f2946, f1094])).
fof(f3939, plain, (~ (e22 = op2(e23, e20)) | ~ spl72_447), inference(forward_demodulation, [], [f2721, f2947])).
fof(f2721, plain, ~ (op2(e23, e20) = h1(e12)), inference(backward_demodulation, [], [f314, f652])).
fof(f314, plain, ~ (op2(e20, e20) = op2(e23, e20)), inference(cnf_transformation, [], [f6])).
fof(f3907, plain, (~ spl72_17 | ~ spl72_49), inference(avatar_split_clause, [], [f3906, f935, f799])).
fof(f799, plain, (spl72_17 <=> (e10 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl72_17])])).
fof(f3906, plain, (~ (e10 = op1(e12, e13)) | ~ spl72_49), inference(forward_demodulation, [], [f283, f937])).
fof(f283, plain, ~ (op1(e10, e13) = op1(e12, e13)), inference(cnf_transformation, [], [f5])).
fof(f3888, plain, (~ spl72_9 | ~ spl72_11), inference(avatar_contradiction_clause, [], [f3887])).
fof(f3887, plain, ($false | (~ spl72_9 | ~ spl72_11)), inference(subsumption_resolution, [], [f3886, f361])).
fof(f3886, plain, ((e10 = e12) | (~ spl72_9 | ~ spl72_11)), inference(backward_demodulation, [], [f775, f767])).
fof(f3857, plain, (~ spl72_50 | ~ spl72_54), inference(avatar_split_clause, [], [f3854, f956, f939])).
fof(f939, plain, (spl72_50 <=> (e11 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl72_50])])).
fof(f3854, plain, (~ (e11 = op1(e10, e13)) | ~ spl72_54), inference(backward_demodulation, [], [f293, f958])).
fof(f293, plain, ~ (op1(e10, e12) = op1(e10, e13)), inference(cnf_transformation, [], [f5])).
fof(f3856, plain, (~ spl72_22 | ~ spl72_54), inference(avatar_split_clause, [], [f3853, f956, f820])).
fof(f820, plain, (spl72_22 <=> (e11 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl72_22])])).
fof(f3853, plain, (~ (e11 = op1(e12, e12)) | ~ spl72_54), inference(backward_demodulation, [], [f277, f958])).
fof(f277, plain, ~ (op1(e10, e12) = op1(e12, e12)), inference(cnf_transformation, [], [f5])).
fof(f3855, plain, (~ spl72_38 | ~ spl72_54), inference(avatar_split_clause, [], [f3852, f956, f888])).
fof(f888, plain, (spl72_38 <=> (e11 = op1(e11, e12))), introduced(avatar_definition, [new_symbols(naming, [spl72_38])])).
fof(f3852, plain, (~ (e11 = op1(e11, e12)) | ~ spl72_54), inference(backward_demodulation, [], [f276, f958])).
fof(f276, plain, ~ (op1(e10, e12) = op1(e11, e12)), inference(cnf_transformation, [], [f5])).
fof(f3849, plain, (~ spl72_66 | spl72_415), inference(avatar_contradiction_clause, [], [f3848])).
fof(f3848, plain, ($false | (~ spl72_66 | spl72_415)), inference(subsumption_resolution, [], [f3847, f2786])).
fof(f2786, plain, (~ (e21 = h4(e12)) | spl72_415), inference(avatar_component_clause, [], [f2784])).
fof(f3847, plain, ((e21 = h4(e12)) | ~ spl72_66), inference(backward_demodulation, [], [f664, f1041])).
fof(f3844, plain, (~ spl72_73 | ~ spl72_75), inference(avatar_contradiction_clause, [], [f3843])).
fof(f3843, plain, ($false | (~ spl72_73 | ~ spl72_75)), inference(subsumption_resolution, [], [f3842, f367])).
fof(f3842, plain, ((e20 = e22) | (~ spl72_73 | ~ spl72_75)), inference(backward_demodulation, [], [f1079, f1071])).
fof(f3834, plain, (spl72_490 | ~ spl72_88), inference(avatar_split_clause, [], [f3833, f1132, f3197])).
fof(f3833, plain, ((e23 = h3(e12)) | ~ spl72_88), inference(backward_demodulation, [], [f660, f1134])).
fof(f3820, plain, (spl72_439 | ~ spl72_106), inference(avatar_split_clause, [], [f3819, f1209, f2905])).
fof(f3819, plain, ((e21 = h2(e12)) | ~ spl72_106), inference(backward_demodulation, [], [f656, f1211])).
fof(f3818, plain, (~ spl72_117 | ~ spl72_447), inference(avatar_contradiction_clause, [], [f3817])).
fof(f3817, plain, ($false | (~ spl72_117 | ~ spl72_447)), inference(subsumption_resolution, [], [f3816, f366])).
fof(f3816, plain, ((e20 = e21) | (~ spl72_117 | ~ spl72_447)), inference(forward_demodulation, [], [f3812, f1258])).
fof(f1258, plain, ((e20 = op2(e20, e22)) | ~ spl72_117), inference(avatar_component_clause, [], [f1256])).
fof(f3815, plain, (~ spl72_95 | ~ spl72_447), inference(avatar_split_clause, [], [f3811, f2946, f1162])).
fof(f3811, plain, (~ (e22 = op2(e22, e20)) | ~ spl72_447), inference(backward_demodulation, [], [f2720, f2947])).
fof(f2720, plain, ~ (op2(e22, e20) = h1(e12)), inference(backward_demodulation, [], [f313, f652])).
fof(f313, plain, ~ (op2(e20, e20) = op2(e22, e20)), inference(cnf_transformation, [], [f6])).
fof(f3808, plain, (~ spl72_92 | ~ spl72_515), inference(avatar_split_clause, [], [f3805, f3362, f1149])).
fof(f3805, plain, (~ (e23 = op2(e22, e21)) | ~ spl72_515), inference(backward_demodulation, [], [f2714, f3363])).
fof(f2714, plain, ~ (op2(e22, e21) = h1(e13)), inference(backward_demodulation, [], [f319, f2712])).
fof(f319, plain, ~ (op2(e20, e21) = op2(e22, e21)), inference(cnf_transformation, [], [f6])).
fof(f3807, plain, (~ spl72_108 | ~ spl72_515), inference(avatar_split_clause, [], [f3804, f3362, f1217])).
fof(f1217, plain, (spl72_108 <=> (e23 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl72_108])])).
fof(f3804, plain, (~ (e23 = op2(e21, e21)) | ~ spl72_515), inference(backward_demodulation, [], [f2713, f3363])).
fof(f2713, plain, ~ (op2(e21, e21) = h1(e13)), inference(backward_demodulation, [], [f318, f2712])).
fof(f318, plain, ~ (op2(e20, e21) = op2(e21, e21)), inference(cnf_transformation, [], [f6])).
fof(f3793, plain, (~ spl72_91 | ~ spl72_83), inference(avatar_split_clause, [], [f3792, f1111, f1145])).
fof(f3792, plain, (~ (e22 = op2(e22, e21)) | ~ spl72_83), inference(forward_demodulation, [], [f352, f1113])).
fof(f3791, plain, (~ spl72_411 | ~ spl72_83), inference(avatar_split_clause, [], [f3790, f1111, f2764])).
fof(f3790, plain, (~ (e22 = h4(e12)) | ~ spl72_83), inference(forward_demodulation, [], [f2747, f1113])).
fof(f3789, plain, (~ spl72_423 | ~ spl72_83), inference(avatar_split_clause, [], [f3788, f1111, f2824])).
fof(f3788, plain, (~ (e22 = h3(e12)) | ~ spl72_83), inference(forward_demodulation, [], [f2741, f1113])).
fof(f3732, plain, (spl72_51 | ~ spl72_5 | ~ spl72_212), inference(avatar_split_clause, [], [f3731, f1744, f748, f943])).
fof(f943, plain, (spl72_51 <=> (e12 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl72_51])])).
fof(f3731, plain, ((e12 = op1(e10, e13)) | (~ spl72_5 | ~ spl72_212)), inference(forward_demodulation, [], [f1746, f750])).
fof(f3709, plain, (~ spl72_8 | spl72_211), inference(avatar_contradiction_clause, [], [f3708])).
fof(f3708, plain, ($false | (~ spl72_8 | spl72_211)), inference(subsumption_resolution, [], [f3704, f762])).
fof(f3704, plain, (~ (e13 = op1(e13, e12)) | (~ spl72_8 | spl72_211)), inference(backward_demodulation, [], [f1741, f762])).
fof(f1741, plain, (~ (e13 = op1(op1(e13, e12), e12)) | spl72_211), inference(avatar_component_clause, [], [f1739])).
fof(f3702, plain, (~ spl72_10 | ~ spl72_11), inference(avatar_contradiction_clause, [], [f3701])).
fof(f3701, plain, ($false | (~ spl72_10 | ~ spl72_11)), inference(subsumption_resolution, [], [f3700, f363])).
fof(f3700, plain, ((e11 = e12) | (~ spl72_10 | ~ spl72_11)), inference(backward_demodulation, [], [f775, f771])).
fof(f3695, plain, (~ spl72_14 | ~ spl72_16), inference(avatar_contradiction_clause, [], [f3694])).
fof(f3694, plain, ($false | (~ spl72_14 | ~ spl72_16)), inference(subsumption_resolution, [], [f3693, f364])).
fof(f3693, plain, ((e11 = e13) | (~ spl72_14 | ~ spl72_16)), inference(backward_demodulation, [], [f796, f788])).
fof(f3692, plain, (~ spl72_4 | ~ spl72_16), inference(avatar_split_clause, [], [f3689, f794, f743])).
fof(f3689, plain, (~ (e13 = op1(e13, e13)) | ~ spl72_16), inference(backward_demodulation, [], [f308, f796])).
fof(f3683, plain, (~ spl72_18 | ~ spl72_19), inference(avatar_contradiction_clause, [], [f3682])).
fof(f3682, plain, ($false | (~ spl72_18 | ~ spl72_19)), inference(subsumption_resolution, [], [f3681, f363])).
fof(f3681, plain, ((e11 = e12) | (~ spl72_18 | ~ spl72_19)), inference(backward_demodulation, [], [f809, f805])).
fof(f3680, plain, (~ spl72_19 | ~ spl72_20), inference(avatar_contradiction_clause, [], [f3679])).
fof(f3679, plain, ($false | (~ spl72_19 | ~ spl72_20)), inference(subsumption_resolution, [], [f3678, f365])).
fof(f3678, plain, ((e12 = e13) | (~ spl72_19 | ~ spl72_20)), inference(backward_demodulation, [], [f813, f809])).
fof(f3645, plain, (~ spl72_33 | ~ spl72_37), inference(avatar_split_clause, [], [f3641, f884, f867])).
fof(f3641, plain, (~ (e10 = op1(e11, e13)) | ~ spl72_37), inference(backward_demodulation, [], [f299, f886])).
fof(f3628, plain, (~ spl72_33 | ~ spl72_45), inference(avatar_split_clause, [], [f3623, f918, f867])).
fof(f3623, plain, (~ (e10 = op1(e11, e13)) | ~ spl72_45), inference(backward_demodulation, [], [f296, f920])).
fof(f3612, plain, (~ spl72_49 | ~ spl72_53), inference(avatar_split_clause, [], [f3608, f952, f935])).
fof(f952, plain, (spl72_53 <=> (e10 = op1(e10, e12))), introduced(avatar_definition, [new_symbols(naming, [spl72_53])])).
fof(f3608, plain, (~ (e10 = op1(e10, e13)) | ~ spl72_53), inference(backward_demodulation, [], [f293, f954])).
fof(f954, plain, ((e10 = op1(e10, e12)) | ~ spl72_53), inference(avatar_component_clause, [], [f952])).
fof(f3604, plain, (~ spl72_52 | ~ spl72_60), inference(avatar_split_clause, [], [f3599, f981, f947])).
fof(f947, plain, (spl72_52 <=> (e13 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl72_52])])).
fof(f3599, plain, (~ (e13 = op1(e10, e13)) | ~ spl72_60), inference(backward_demodulation, [], [f292, f983])).
fof(f292, plain, ~ (op1(e10, e11) = op1(e10, e13)), inference(cnf_transformation, [], [f5])).
fof(f3603, plain, (~ spl72_56 | ~ spl72_60), inference(avatar_split_clause, [], [f3598, f981, f964])).
fof(f964, plain, (spl72_56 <=> (e13 = op1(e10, e12))), introduced(avatar_definition, [new_symbols(naming, [spl72_56])])).
fof(f3598, plain, (~ (e13 = op1(e10, e12)) | ~ spl72_60), inference(backward_demodulation, [], [f291, f983])).
fof(f291, plain, ~ (op1(e10, e11) = op1(e10, e12)), inference(cnf_transformation, [], [f5])).
fof(f3602, plain, (~ spl72_12 | ~ spl72_60), inference(avatar_split_clause, [], [f3597, f981, f777])).
fof(f777, plain, (spl72_12 <=> (e13 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl72_12])])).
fof(f3597, plain, (~ (e13 = op1(e13, e11)) | ~ spl72_60), inference(backward_demodulation, [], [f272, f983])).
fof(f272, plain, ~ (op1(e10, e11) = op1(e13, e11)), inference(cnf_transformation, [], [f5])).
fof(f3601, plain, (~ spl72_28 | ~ spl72_60), inference(avatar_split_clause, [], [f3596, f981, f845])).
fof(f845, plain, (spl72_28 <=> (e13 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl72_28])])).
fof(f3596, plain, (~ (e13 = op1(e12, e11)) | ~ spl72_60), inference(backward_demodulation, [], [f271, f983])).
fof(f271, plain, ~ (op1(e10, e11) = op1(e12, e11)), inference(cnf_transformation, [], [f5])).
fof(f3600, plain, (~ spl72_44 | ~ spl72_60), inference(avatar_split_clause, [], [f3595, f981, f913])).
fof(f913, plain, (spl72_44 <=> (e13 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl72_44])])).
fof(f3595, plain, (~ (e13 = op1(e11, e11)) | ~ spl72_60), inference(backward_demodulation, [], [f270, f983])).
fof(f270, plain, ~ (op1(e10, e11) = op1(e11, e11)), inference(cnf_transformation, [], [f5])).
fof(f3594, plain, (spl72_54 | ~ spl72_63), inference(avatar_split_clause, [], [f3584, f994, f956])).
fof(f3584, plain, ((e11 = op1(e10, e12)) | ~ spl72_63), inference(backward_demodulation, [], [f644, f996])).
fof(f644, plain, (e11 = op1(e10, op1(e10, e10))), inference(cnf_transformation, [], [f12])).
fof(f12, plain, ((e13 = op1(e10, op1(e10, op1(e10, e10)))) & (op1(e10, e10) = e12) & (e11 = op1(e10, op1(e10, e10)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG104+1.p', ax12)).
fof(f3593, plain, (~ spl72_51 | ~ spl72_63), inference(avatar_split_clause, [], [f3583, f994, f943])).
fof(f3583, plain, (~ (e12 = op1(e10, e13)) | ~ spl72_63), inference(backward_demodulation, [], [f290, f996])).
fof(f290, plain, ~ (op1(e10, e10) = op1(e10, e13)), inference(cnf_transformation, [], [f5])).
fof(f3592, plain, (~ spl72_55 | ~ spl72_63), inference(avatar_split_clause, [], [f3582, f994, f960])).
fof(f960, plain, (spl72_55 <=> (e12 = op1(e10, e12))), introduced(avatar_definition, [new_symbols(naming, [spl72_55])])).
fof(f3582, plain, (~ (e12 = op1(e10, e12)) | ~ spl72_63), inference(backward_demodulation, [], [f289, f996])).
fof(f289, plain, ~ (op1(e10, e10) = op1(e10, e12)), inference(cnf_transformation, [], [f5])).
fof(f3591, plain, (~ spl72_59 | ~ spl72_63), inference(avatar_split_clause, [], [f3581, f994, f977])).
fof(f977, plain, (spl72_59 <=> (e12 = op1(e10, e11))), introduced(avatar_definition, [new_symbols(naming, [spl72_59])])).
fof(f3581, plain, (~ (e12 = op1(e10, e11)) | ~ spl72_63), inference(backward_demodulation, [], [f288, f996])).
fof(f288, plain, ~ (op1(e10, e10) = op1(e10, e11)), inference(cnf_transformation, [], [f5])).
fof(f3590, plain, (~ spl72_15 | ~ spl72_63), inference(avatar_split_clause, [], [f3580, f994, f790])).
fof(f790, plain, (spl72_15 <=> (e12 = op1(e13, e10))), introduced(avatar_definition, [new_symbols(naming, [spl72_15])])).
fof(f3580, plain, (~ (e12 = op1(e13, e10)) | ~ spl72_63), inference(backward_demodulation, [], [f266, f996])).
fof(f266, plain, ~ (op1(e10, e10) = op1(e13, e10)), inference(cnf_transformation, [], [f5])).
fof(f3589, plain, (~ spl72_31 | ~ spl72_63), inference(avatar_split_clause, [], [f3579, f994, f858])).
fof(f858, plain, (spl72_31 <=> (e12 = op1(e12, e10))), introduced(avatar_definition, [new_symbols(naming, [spl72_31])])).
fof(f3579, plain, (~ (e12 = op1(e12, e10)) | ~ spl72_63), inference(backward_demodulation, [], [f265, f996])).
fof(f265, plain, ~ (op1(e10, e10) = op1(e12, e10)), inference(cnf_transformation, [], [f5])).
fof(f3588, plain, (~ spl72_47 | ~ spl72_63), inference(avatar_split_clause, [], [f3578, f994, f926])).
fof(f3578, plain, (~ (e12 = op1(e11, e10)) | ~ spl72_63), inference(backward_demodulation, [], [f264, f996])).
fof(f264, plain, ~ (op1(e10, e10) = op1(e11, e10)), inference(cnf_transformation, [], [f5])).
fof(f3577, plain, (spl72_419 | ~ spl72_65), inference(avatar_split_clause, [], [f3576, f1035, f2804])).
fof(f3576, plain, ((e20 = h4(e12)) | ~ spl72_65), inference(backward_demodulation, [], [f664, f1037])).
fof(f1037, plain, ((e20 = op2(e23, e23)) | ~ spl72_65), inference(avatar_component_clause, [], [f1035])).
fof(f3572, plain, (~ spl72_71 | ~ spl72_72), inference(avatar_contradiction_clause, [], [f3571])).
fof(f3571, plain, ($false | (~ spl72_71 | ~ spl72_72)), inference(subsumption_resolution, [], [f3570, f371])).
fof(f3570, plain, ((e22 = e23) | (~ spl72_71 | ~ spl72_72)), inference(backward_demodulation, [], [f1066, f1062])).
fof(f1066, plain, ((e23 = op2(e23, e22)) | ~ spl72_72), inference(avatar_component_clause, [], [f1064])).
fof(f3567, plain, (~ spl72_67 | ~ spl72_72 | spl72_353), inference(avatar_split_clause, [], [f3561, f2433, f1064, f1043])).
fof(f3561, plain, (~ (e22 = op2(e23, e23)) | (~ spl72_72 | spl72_353)), inference(backward_demodulation, [], [f2435, f1066])).
fof(f3566, plain, (spl72_67 | ~ spl72_72 | ~ spl72_352), inference(avatar_split_clause, [], [f3560, f2428, f1064, f1043])).
fof(f3560, plain, ((e22 = op2(e23, e23)) | (~ spl72_72 | ~ spl72_352)), inference(backward_demodulation, [], [f2430, f1066])).
fof(f3565, plain, (~ spl72_72 | spl72_351), inference(avatar_contradiction_clause, [], [f3564])).
fof(f3564, plain, ($false | (~ spl72_72 | spl72_351)), inference(subsumption_resolution, [], [f3559, f1066])).
fof(f3559, plain, (~ (e23 = op2(e23, e22)) | (~ spl72_72 | spl72_351)), inference(backward_demodulation, [], [f2425, f1066])).
fof(f3547, plain, (~ spl72_78 | ~ spl72_80), inference(avatar_contradiction_clause, [], [f3546])).
fof(f3546, plain, ($false | (~ spl72_78 | ~ spl72_80)), inference(subsumption_resolution, [], [f3545, f370])).
fof(f3545, plain, ((e21 = e23) | (~ spl72_78 | ~ spl72_80)), inference(backward_demodulation, [], [f1100, f1092])).
fof(f3544, plain, (~ spl72_72 | ~ spl72_80), inference(avatar_split_clause, [], [f3540, f1098, f1064])).
fof(f3540, plain, (~ (e23 = op2(e23, e22)) | ~ spl72_80), inference(backward_demodulation, [], [f355, f1100])).
fof(f3536, plain, (~ spl72_82 | ~ spl72_83), inference(avatar_contradiction_clause, [], [f3535])).
fof(f3535, plain, ($false | (~ spl72_82 | ~ spl72_83)), inference(subsumption_resolution, [], [f3534, f369])).
fof(f369, plain, ~ (e21 = e22), inference(cnf_transformation, [], [f8])).
fof(f3534, plain, ((e21 = e22) | (~ spl72_82 | ~ spl72_83)), inference(backward_demodulation, [], [f1113, f1109])).
fof(f3526, plain, (spl72_431 | ~ spl72_85), inference(avatar_split_clause, [], [f3525, f1120, f2865])).
fof(f1120, plain, (spl72_85 <=> (e20 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl72_85])])).
fof(f3525, plain, ((e20 = h3(e12)) | ~ spl72_85), inference(backward_demodulation, [], [f660, f1122])).
fof(f1122, plain, ((e20 = op2(e22, e22)) | ~ spl72_85), inference(avatar_component_clause, [], [f1120])).
fof(f3501, plain, (~ spl72_99 | ~ spl72_100), inference(avatar_contradiction_clause, [], [f3500])).
fof(f3500, plain, ($false | (~ spl72_99 | ~ spl72_100)), inference(subsumption_resolution, [], [f3499, f371])).
fof(f3499, plain, ((e22 = e23) | (~ spl72_99 | ~ spl72_100)), inference(backward_demodulation, [], [f1185, f1181])).
fof(f3490, plain, (~ spl72_443 | ~ spl72_101), inference(avatar_split_clause, [], [f3486, f1188, f2926])).
fof(f3486, plain, (~ (e20 = h2(e12)) | ~ spl72_101), inference(backward_demodulation, [], [f2731, f1190])).
fof(f3483, plain, (spl72_443 | ~ spl72_105), inference(avatar_split_clause, [], [f3481, f1205, f2926])).
fof(f1205, plain, (spl72_105 <=> (e20 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl72_105])])).
fof(f3481, plain, ((e20 = h2(e12)) | ~ spl72_105), inference(backward_demodulation, [], [f656, f1207])).
fof(f1207, plain, ((e20 = op2(e21, e21)) | ~ spl72_105), inference(avatar_component_clause, [], [f1205])).
fof(f3446, plain, (spl72_515 | ~ spl72_124), inference(avatar_split_clause, [], [f3445, f1285, f3362])).
fof(f3445, plain, ((e23 = h1(e13)) | ~ spl72_124), inference(backward_demodulation, [], [f2712, f1287])).
fof(f3444, plain, (spl72_447 | ~ spl72_127), inference(avatar_split_clause, [], [f3443, f1298, f2946])).
fof(f3443, plain, ((e22 = h1(e12)) | ~ spl72_127), inference(backward_demodulation, [], [f652, f1300])).
fof(f3413, plain, (~ spl72_511 | ~ spl72_512 | ~ spl72_513 | ~ spl72_514 | spl72_451 | spl72_448 | spl72_445 | ~ spl72_515 | ~ spl72_516 | ~ spl72_517 | ~ spl72_518 | ~ spl72_519 | ~ spl72_520 | ~ spl72_521 | ~ spl72_522 | ~ spl72_523 | ~ spl72_524 | ~ spl72_525 | ~ spl72_526 | ~ spl72_527), inference(avatar_split_clause, [], [f3344, f3410, f3406, f3402, f3398, f3394, f3390, f3386, f3382, f3378, f3374, f3370, f3366, f3362, f2937, f2953, f2970, f3358, f3354, f3350, f3346])).
fof(f2970, plain, (spl72_451 <=> sP60), introduced(avatar_definition, [new_symbols(naming, [spl72_451])])).
fof(f2953, plain, (spl72_448 <=> sP61), introduced(avatar_definition, [new_symbols(naming, [spl72_448])])).
fof(f2937, plain, (spl72_445 <=> sP62), introduced(avatar_definition, [new_symbols(naming, [spl72_445])])).
fof(f3344, plain, (~ (h1(e12) = h1(op1(e10, e10))) | ~ (h1(e13) = h1(op1(e10, e11))) | ~ (e21 = h1(op1(e10, e12))) | ~ (h1(op1(e10, e13)) = op2(e20, h1(e13))) | ~ (op2(e21, e20) = h1(op1(e11, e10))) | ~ (h2(e12) = h1(op1(e11, e11))) | ~ (h1(op1(e11, e12)) = op2(e21, h1(e12))) | ~ (h1(op1(e11, e13)) = op2(e21, h1(e13))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), e20)) | ~ (h1(op1(e12, e11)) = op2(h1(e12), e21)) | ~ (h1(op1(e13, e10)) = op2(h1(e13), e20)) | ~ (h1(op1(e13, e11)) = op2(h1(e13), e21)) | ~ (e23 = h1(e13)) | sP62 | sP61 | sP60 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12)))), inference(forward_demodulation, [], [f3343, f652])).
fof(f3343, plain, (~ (op2(e20, e20) = h1(op1(e10, e10))) | ~ (h1(e13) = h1(op1(e10, e11))) | ~ (e21 = h1(op1(e10, e12))) | ~ (h1(op1(e10, e13)) = op2(e20, h1(e13))) | ~ (op2(e21, e20) = h1(op1(e11, e10))) | ~ (h2(e12) = h1(op1(e11, e11))) | ~ (h1(op1(e11, e12)) = op2(e21, h1(e12))) | ~ (h1(op1(e11, e13)) = op2(e21, h1(e13))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), e20)) | ~ (h1(op1(e12, e11)) = op2(h1(e12), e21)) | ~ (h1(op1(e13, e10)) = op2(h1(e13), e20)) | ~ (h1(op1(e13, e11)) = op2(h1(e13), e21)) | ~ (e23 = h1(e13)) | sP62 | sP61 | sP60 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12)))), inference(forward_demodulation, [], [f3342, f650])).
fof(f3342, plain, (~ (h1(e13) = h1(op1(e10, e11))) | ~ (e21 = h1(op1(e10, e12))) | ~ (h1(op1(e10, e13)) = op2(e20, h1(e13))) | ~ (op2(e21, e20) = h1(op1(e11, e10))) | ~ (h2(e12) = h1(op1(e11, e11))) | ~ (h1(op1(e11, e12)) = op2(e21, h1(e12))) | ~ (h1(op1(e11, e13)) = op2(e21, h1(e13))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), e20)) | ~ (h1(op1(e12, e11)) = op2(h1(e12), e21)) | ~ (h1(op1(e13, e10)) = op2(h1(e13), e20)) | ~ (h1(op1(e13, e11)) = op2(h1(e13), e21)) | ~ (e23 = h1(e13)) | sP62 | sP61 | sP60 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f3341, f2712])).
fof(f3341, plain, (~ (op2(e20, e21) = h1(op1(e10, e11))) | ~ (e21 = h1(op1(e10, e12))) | ~ (h1(op1(e10, e13)) = op2(e20, h1(e13))) | ~ (op2(e21, e20) = h1(op1(e11, e10))) | ~ (h2(e12) = h1(op1(e11, e11))) | ~ (h1(op1(e11, e12)) = op2(e21, h1(e12))) | ~ (h1(op1(e11, e13)) = op2(e21, h1(e13))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), e20)) | ~ (h1(op1(e12, e11)) = op2(h1(e12), e21)) | ~ (h1(op1(e13, e10)) = op2(h1(e13), e20)) | ~ (h1(op1(e13, e11)) = op2(h1(e13), e21)) | ~ (e23 = h1(e13)) | sP62 | sP61 | sP60 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f3340, f650])).
fof(f3340, plain, (~ (h1(op1(e10, e11)) = op2(h1(e10), e21)) | ~ (e21 = h1(op1(e10, e12))) | ~ (h1(op1(e10, e13)) = op2(e20, h1(e13))) | ~ (op2(e21, e20) = h1(op1(e11, e10))) | ~ (h2(e12) = h1(op1(e11, e11))) | ~ (h1(op1(e11, e12)) = op2(e21, h1(e12))) | ~ (h1(op1(e11, e13)) = op2(e21, h1(e13))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), e20)) | ~ (h1(op1(e12, e11)) = op2(h1(e12), e21)) | ~ (h1(op1(e13, e10)) = op2(h1(e13), e20)) | ~ (h1(op1(e13, e11)) = op2(h1(e13), e21)) | ~ (e23 = h1(e13)) | sP62 | sP61 | sP60 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f3339, f2727])).
fof(f3339, plain, (~ (e21 = h1(op1(e10, e12))) | ~ (h1(op1(e10, e13)) = op2(e20, h1(e13))) | ~ (op2(e21, e20) = h1(op1(e11, e10))) | ~ (h2(e12) = h1(op1(e11, e11))) | ~ (h1(op1(e11, e12)) = op2(e21, h1(e12))) | ~ (h1(op1(e11, e13)) = op2(e21, h1(e13))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), e20)) | ~ (h1(op1(e12, e11)) = op2(h1(e12), e21)) | ~ (h1(op1(e13, e10)) = op2(h1(e13), e20)) | ~ (h1(op1(e13, e11)) = op2(h1(e13), e21)) | ~ (e23 = h1(e13)) | sP62 | sP61 | sP60 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f3338, f2724])).
fof(f3338, plain, (~ (h1(op1(e10, e12)) = op2(e20, h1(e12))) | ~ (h1(op1(e10, e13)) = op2(e20, h1(e13))) | ~ (op2(e21, e20) = h1(op1(e11, e10))) | ~ (h2(e12) = h1(op1(e11, e11))) | ~ (h1(op1(e11, e12)) = op2(e21, h1(e12))) | ~ (h1(op1(e11, e13)) = op2(e21, h1(e13))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), e20)) | ~ (h1(op1(e12, e11)) = op2(h1(e12), e21)) | ~ (h1(op1(e13, e10)) = op2(h1(e13), e20)) | ~ (h1(op1(e13, e11)) = op2(h1(e13), e21)) | ~ (e23 = h1(e13)) | sP62 | sP61 | sP60 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f3337, f650])).
fof(f3337, plain, (~ (h1(op1(e10, e13)) = op2(e20, h1(e13))) | ~ (op2(e21, e20) = h1(op1(e11, e10))) | ~ (h2(e12) = h1(op1(e11, e11))) | ~ (h1(op1(e11, e12)) = op2(e21, h1(e12))) | ~ (h1(op1(e11, e13)) = op2(e21, h1(e13))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), e20)) | ~ (h1(op1(e12, e11)) = op2(h1(e12), e21)) | ~ (h1(op1(e13, e10)) = op2(h1(e13), e20)) | ~ (h1(op1(e13, e11)) = op2(h1(e13), e21)) | ~ (e23 = h1(e13)) | sP62 | sP61 | sP60 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f3336, f650])).
fof(f3336, plain, (~ (op2(e21, e20) = h1(op1(e11, e10))) | ~ (h2(e12) = h1(op1(e11, e11))) | ~ (h1(op1(e11, e12)) = op2(e21, h1(e12))) | ~ (h1(op1(e11, e13)) = op2(e21, h1(e13))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), e20)) | ~ (h1(op1(e12, e11)) = op2(h1(e12), e21)) | ~ (h1(op1(e13, e10)) = op2(h1(e13), e20)) | ~ (h1(op1(e13, e11)) = op2(h1(e13), e21)) | ~ (e23 = h1(e13)) | sP62 | sP61 | sP60 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f3335, f2727])).
fof(f3335, plain, (~ (h1(op1(e11, e10)) = op2(h1(e11), e20)) | ~ (h2(e12) = h1(op1(e11, e11))) | ~ (h1(op1(e11, e12)) = op2(e21, h1(e12))) | ~ (h1(op1(e11, e13)) = op2(e21, h1(e13))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), e20)) | ~ (h1(op1(e12, e11)) = op2(h1(e12), e21)) | ~ (h1(op1(e13, e10)) = op2(h1(e13), e20)) | ~ (h1(op1(e13, e11)) = op2(h1(e13), e21)) | ~ (e23 = h1(e13)) | sP62 | sP61 | sP60 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f3334, f650])).
fof(f3334, plain, (~ (h2(e12) = h1(op1(e11, e11))) | ~ (h1(op1(e11, e12)) = op2(e21, h1(e12))) | ~ (h1(op1(e11, e13)) = op2(e21, h1(e13))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), e20)) | ~ (h1(op1(e12, e11)) = op2(h1(e12), e21)) | ~ (h1(op1(e13, e10)) = op2(h1(e13), e20)) | ~ (h1(op1(e13, e11)) = op2(h1(e13), e21)) | ~ (e23 = h1(e13)) | sP62 | sP61 | sP60 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f3333, f656])).
fof(f3333, plain, (~ (op2(e21, e21) = h1(op1(e11, e11))) | ~ (h1(op1(e11, e12)) = op2(e21, h1(e12))) | ~ (h1(op1(e11, e13)) = op2(e21, h1(e13))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), e20)) | ~ (h1(op1(e12, e11)) = op2(h1(e12), e21)) | ~ (h1(op1(e13, e10)) = op2(h1(e13), e20)) | ~ (h1(op1(e13, e11)) = op2(h1(e13), e21)) | ~ (e23 = h1(e13)) | sP62 | sP61 | sP60 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f3332, f2727])).
fof(f3332, plain, (~ (h1(op1(e11, e12)) = op2(e21, h1(e12))) | ~ (h1(op1(e11, e13)) = op2(e21, h1(e13))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), e20)) | ~ (h1(op1(e12, e11)) = op2(h1(e12), e21)) | ~ (h1(op1(e13, e10)) = op2(h1(e13), e20)) | ~ (h1(op1(e13, e11)) = op2(h1(e13), e21)) | ~ (e23 = h1(e13)) | sP62 | sP61 | sP60 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f3331, f2727])).
fof(f3331, plain, (~ (h1(op1(e11, e13)) = op2(e21, h1(e13))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), e20)) | ~ (h1(op1(e12, e11)) = op2(h1(e12), e21)) | ~ (h1(op1(e13, e10)) = op2(h1(e13), e20)) | ~ (h1(op1(e13, e11)) = op2(h1(e13), e21)) | ~ (e23 = h1(e13)) | sP62 | sP61 | sP60 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f3330, f2727])).
fof(f3330, plain, (~ (h1(op1(e12, e10)) = op2(h1(e12), e20)) | ~ (h1(op1(e12, e11)) = op2(h1(e12), e21)) | ~ (h1(op1(e13, e10)) = op2(h1(e13), e20)) | ~ (h1(op1(e13, e11)) = op2(h1(e13), e21)) | ~ (e23 = h1(e13)) | sP62 | sP61 | sP60 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f3329, f650])).
fof(f3329, plain, (~ (h1(op1(e12, e11)) = op2(h1(e12), e21)) | ~ (h1(op1(e13, e10)) = op2(h1(e13), e20)) | ~ (h1(op1(e13, e11)) = op2(h1(e13), e21)) | ~ (e23 = h1(e13)) | sP62 | sP61 | sP60 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f3328, f2727])).
fof(f3328, plain, (~ (h1(op1(e13, e10)) = op2(h1(e13), e20)) | ~ (h1(op1(e13, e11)) = op2(h1(e13), e21)) | ~ (e23 = h1(e13)) | sP62 | sP61 | sP60 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f3327, f650])).
fof(f3327, plain, (~ (h1(op1(e13, e11)) = op2(h1(e13), e21)) | ~ (e23 = h1(e13)) | sP62 | sP61 | sP60 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f717, f2727])).
fof(f717, plain, (~ (e23 = h1(e13)) | sP62 | sP61 | sP60 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(cnf_transformation, [], [f95])).
fof(f95, plain, (((~ (e23 = h4(e13)) & ~ (e23 = h4(e12)) & ~ (e23 = h4(e11)) & ~ (e23 = h4(e10))) | sP71 | sP70 | sP69 | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) & ((~ (e23 = h3(e13)) & ~ (e23 = h3(e12)) & ~ (e23 = h3(e11)) & ~ (e23 = h3(e10))) | sP68 | sP67 | sP66 | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) & ((~ (e23 = h2(e13)) & ~ (e23 = h2(e12)) & ~ (e23 = h2(e11)) & ~ (e23 = h2(e10))) | sP65 | sP64 | sP63 | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) & ((~ (e23 = h1(e13)) & ~ (e23 = h1(e12)) & ~ (e23 = h1(e11)) & ~ (e23 = h1(e10))) | sP62 | sP61 | sP60 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(definition_folding, [], [f20, e94, e93, e92, e91, e90, e89, e88, e87, e86, e85, e84, e83])).
fof(f83, plain, ((~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10))) | ~ sP60), inference(usedef, [], [e83])).
fof(e83, plain, (sP60 <=> (~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP60])])).
fof(f84, plain, ((~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10))) | ~ sP61), inference(usedef, [], [e84])).
fof(e84, plain, (sP61 <=> (~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP61])])).
fof(f85, plain, ((~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10))) | ~ sP62), inference(usedef, [], [e85])).
fof(e85, plain, (sP62 <=> (~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP62])])).
fof(f86, plain, ((~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10))) | ~ sP63), inference(usedef, [], [e86])).
fof(e86, plain, (sP63 <=> (~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP63])])).
fof(f87, plain, ((~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10))) | ~ sP64), inference(usedef, [], [e87])).
fof(e87, plain, (sP64 <=> (~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP64])])).
fof(f88, plain, ((~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10))) | ~ sP65), inference(usedef, [], [e88])).
fof(e88, plain, (sP65 <=> (~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP65])])).
fof(f89, plain, ((~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10))) | ~ sP66), inference(usedef, [], [e89])).
fof(e89, plain, (sP66 <=> (~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP66])])).
fof(f90, plain, ((~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10))) | ~ sP67), inference(usedef, [], [e90])).
fof(e90, plain, (sP67 <=> (~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP67])])).
fof(f91, plain, ((~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10))) | ~ sP68), inference(usedef, [], [e91])).
fof(e91, plain, (sP68 <=> (~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP68])])).
fof(f92, plain, ((~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10))) | ~ sP69), inference(usedef, [], [e92])).
fof(e92, plain, (sP69 <=> (~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP69])])).
fof(f93, plain, ((~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10))) | ~ sP70), inference(usedef, [], [e93])).
fof(e93, plain, (sP70 <=> (~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP70])])).
fof(f94, plain, ((~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10))) | ~ sP71), inference(usedef, [], [e94])).
fof(e94, plain, (sP71 <=> (~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP71])])).
fof(f20, plain, (((~ (e23 = h4(e13)) & ~ (e23 = h4(e12)) & ~ (e23 = h4(e11)) & ~ (e23 = h4(e10))) | (~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10))) | (~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10))) | (~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10))) | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) & ((~ (e23 = h3(e13)) & ~ (e23 = h3(e12)) & ~ (e23 = h3(e11)) & ~ (e23 = h3(e10))) | (~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10))) | (~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10))) | (~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10))) | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) & ((~ (e23 = h2(e13)) & ~ (e23 = h2(e12)) & ~ (e23 = h2(e11)) & ~ (e23 = h2(e10))) | (~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10))) | (~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10))) | (~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) & ((~ (e23 = h1(e13)) & ~ (e23 = h1(e12)) & ~ (e23 = h1(e11)) & ~ (e23 = h1(e10))) | (~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10))) | (~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10))) | (~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10))) | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(ennf_transformation, [], [f19])).
fof(f19, plain, ~ ((((e23 = h4(e13)) | (e23 = h4(e12)) | (e23 = h4(e11)) | (e23 = h4(e10))) & ((e22 = h4(e13)) | (e22 = h4(e12)) | (e22 = h4(e11)) | (e22 = h4(e10))) & ((e21 = h4(e13)) | (e21 = h4(e12)) | (e21 = h4(e11)) | (e21 = h4(e10))) & ((e20 = h4(e13)) | (e20 = h4(e12)) | (e20 = h4(e11)) | (e20 = h4(e10))) & (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) & (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) & (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) & (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) & (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) & (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) & (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) & (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) & (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) & (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) & (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) & (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) & (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) & (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) & (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) & (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) | (((e23 = h3(e13)) | (e23 = h3(e12)) | (e23 = h3(e11)) | (e23 = h3(e10))) & ((e22 = h3(e13)) | (e22 = h3(e12)) | (e22 = h3(e11)) | (e22 = h3(e10))) & ((e21 = h3(e13)) | (e21 = h3(e12)) | (e21 = h3(e11)) | (e21 = h3(e10))) & ((e20 = h3(e13)) | (e20 = h3(e12)) | (e20 = h3(e11)) | (e20 = h3(e10))) & (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) & (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) & (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) & (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) & (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) & (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) & (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) & (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) & (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) & (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) & (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) & (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) & (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) & (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) & (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) & (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) | (((e23 = h2(e13)) | (e23 = h2(e12)) | (e23 = h2(e11)) | (e23 = h2(e10))) & ((e22 = h2(e13)) | (e22 = h2(e12)) | (e22 = h2(e11)) | (e22 = h2(e10))) & ((e21 = h2(e13)) | (e21 = h2(e12)) | (e21 = h2(e11)) | (e21 = h2(e10))) & ((e20 = h2(e13)) | (e20 = h2(e12)) | (e20 = h2(e11)) | (e20 = h2(e10))) & (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) & (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) & (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) & (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) & (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) & (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) & (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) & (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) & (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) & (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) & (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) & (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) & (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) & (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) & (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) & (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) | (((e23 = h1(e13)) | (e23 = h1(e12)) | (e23 = h1(e11)) | (e23 = h1(e10))) & ((e22 = h1(e13)) | (e22 = h1(e12)) | (e22 = h1(e11)) | (e22 = h1(e10))) & ((e21 = h1(e13)) | (e21 = h1(e12)) | (e21 = h1(e11)) | (e21 = h1(e10))) & ((e20 = h1(e13)) | (e20 = h1(e12)) | (e20 = h1(e11)) | (e20 = h1(e10))) & (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) & (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) & (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) & (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) & (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) & (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) & (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) & (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) & (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) & (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) & (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) & (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) & (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) & (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) & (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) & (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(negated_conjecture, [], [f18])).
fof(f18, plain, ~ ((((e23 = h4(e13)) | (e23 = h4(e12)) | (e23 = h4(e11)) | (e23 = h4(e10))) & ((e22 = h4(e13)) | (e22 = h4(e12)) | (e22 = h4(e11)) | (e22 = h4(e10))) & ((e21 = h4(e13)) | (e21 = h4(e12)) | (e21 = h4(e11)) | (e21 = h4(e10))) & ((e20 = h4(e13)) | (e20 = h4(e12)) | (e20 = h4(e11)) | (e20 = h4(e10))) & (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) & (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) & (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) & (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) & (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) & (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) & (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) & (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) & (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) & (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) & (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) & (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) & (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) & (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) & (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) & (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) | (((e23 = h3(e13)) | (e23 = h3(e12)) | (e23 = h3(e11)) | (e23 = h3(e10))) & ((e22 = h3(e13)) | (e22 = h3(e12)) | (e22 = h3(e11)) | (e22 = h3(e10))) & ((e21 = h3(e13)) | (e21 = h3(e12)) | (e21 = h3(e11)) | (e21 = h3(e10))) & ((e20 = h3(e13)) | (e20 = h3(e12)) | (e20 = h3(e11)) | (e20 = h3(e10))) & (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) & (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) & (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) & (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) & (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) & (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) & (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) & (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) & (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) & (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) & (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) & (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) & (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) & (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) & (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) & (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) | (((e23 = h2(e13)) | (e23 = h2(e12)) | (e23 = h2(e11)) | (e23 = h2(e10))) & ((e22 = h2(e13)) | (e22 = h2(e12)) | (e22 = h2(e11)) | (e22 = h2(e10))) & ((e21 = h2(e13)) | (e21 = h2(e12)) | (e21 = h2(e11)) | (e21 = h2(e10))) & ((e20 = h2(e13)) | (e20 = h2(e12)) | (e20 = h2(e11)) | (e20 = h2(e10))) & (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) & (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) & (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) & (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) & (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) & (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) & (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) & (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) & (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) & (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) & (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) & (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) & (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) & (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) & (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) & (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) | (((e23 = h1(e13)) | (e23 = h1(e12)) | (e23 = h1(e11)) | (e23 = h1(e10))) & ((e22 = h1(e13)) | (e22 = h1(e12)) | (e22 = h1(e11)) | (e22 = h1(e10))) & ((e21 = h1(e13)) | (e21 = h1(e12)) | (e21 = h1(e11)) | (e21 = h1(e10))) & ((e20 = h1(e13)) | (e20 = h1(e12)) | (e20 = h1(e11)) | (e20 = h1(e10))) & (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) & (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) & (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) & (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) & (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) & (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) & (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) & (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) & (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) & (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) & (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) & (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) & (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) & (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) & (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) & (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG104+1.p', co1)).
fof(f2985, plain, ~ spl72_451, inference(avatar_split_clause, [], [f2984, f2970])).
fof(f2984, plain, ~ sP60, inference(subsumption_resolution, [], [f710, f650])).
fof(f710, plain, (~ (e20 = h1(e10)) | ~ sP60), inference(cnf_transformation, [], [f167])).
fof(f167, plain, ((~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10))) | ~ sP60), inference(nnf_transformation, [], [f83])).
fof(f2967, plain, ~ spl72_448, inference(avatar_split_clause, [], [f2966, f2953])).
fof(f2966, plain, ~ sP61, inference(subsumption_resolution, [], [f707, f2727])).
fof(f707, plain, (~ (e21 = h1(e11)) | ~ sP61), inference(cnf_transformation, [], [f166])).
fof(f166, plain, ((~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10))) | ~ sP61), inference(nnf_transformation, [], [f84])).
fof(f2949, plain, (~ spl72_445 | ~ spl72_447), inference(avatar_split_clause, [], [f704, f2946, f2937])).
fof(f704, plain, (~ (e22 = h1(e12)) | ~ sP62), inference(cnf_transformation, [], [f165])).
fof(f165, plain, ((~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10))) | ~ sP62), inference(nnf_transformation, [], [f85])).
fof(f2711, plain, spl72_124, inference(avatar_split_clause, [], [f2710, f1285])).
fof(f2710, plain, (e23 = op2(e20, e21)), inference(backward_demodulation, [], [f649, f647])).
fof(f649, plain, (e23 = op2(e20, op2(e20, op2(e20, e20)))), inference(cnf_transformation, [], [f13])).
fof(f2709, plain, spl72_127, inference(avatar_split_clause, [], [f648, f1298])).
fof(f648, plain, (op2(e20, e20) = e22), inference(cnf_transformation, [], [f13])).
fof(f2708, plain, spl72_60, inference(avatar_split_clause, [], [f2707, f981])).
fof(f2707, plain, (e13 = op1(e10, e11)), inference(backward_demodulation, [], [f646, f644])).
fof(f646, plain, (e13 = op1(e10, op1(e10, op1(e10, e10)))), inference(cnf_transformation, [], [f12])).
fof(f2706, plain, spl72_63, inference(avatar_split_clause, [], [f645, f994])).
fof(f645, plain, (op1(e10, e10) = e12), inference(cnf_transformation, [], [f12])).
fof(f2705, plain, (spl72_402 | spl72_398 | spl72_394 | spl72_390 | spl72_386 | spl72_384 | spl72_380 | spl72_376 | spl72_372 | spl72_368 | spl72_366 | spl72_362 | spl72_358 | spl72_354 | spl72_350 | ~ spl72_404), inference(avatar_split_clause, [], [f636, f2680, f2419, f2438, f2457, f2476, f2495, f2506, f2525, f2544, f2563, f2582, f2593, f2612, f2631, f2650, f2669])).
fof(f2669, plain, (spl72_402 <=> sP30), introduced(avatar_definition, [new_symbols(naming, [spl72_402])])).
fof(f2650, plain, (spl72_398 <=> sP31), introduced(avatar_definition, [new_symbols(naming, [spl72_398])])).
fof(f2631, plain, (spl72_394 <=> sP32), introduced(avatar_definition, [new_symbols(naming, [spl72_394])])).
fof(f2612, plain, (spl72_390 <=> sP33), introduced(avatar_definition, [new_symbols(naming, [spl72_390])])).
fof(f2593, plain, (spl72_386 <=> sP34), introduced(avatar_definition, [new_symbols(naming, [spl72_386])])).
fof(f2582, plain, (spl72_384 <=> sP35), introduced(avatar_definition, [new_symbols(naming, [spl72_384])])).
fof(f2563, plain, (spl72_380 <=> sP36), introduced(avatar_definition, [new_symbols(naming, [spl72_380])])).
fof(f2544, plain, (spl72_376 <=> sP37), introduced(avatar_definition, [new_symbols(naming, [spl72_376])])).
fof(f2525, plain, (spl72_372 <=> sP38), introduced(avatar_definition, [new_symbols(naming, [spl72_372])])).
fof(f2506, plain, (spl72_368 <=> sP39), introduced(avatar_definition, [new_symbols(naming, [spl72_368])])).
fof(f2495, plain, (spl72_366 <=> sP40), introduced(avatar_definition, [new_symbols(naming, [spl72_366])])).
fof(f2476, plain, (spl72_362 <=> sP41), introduced(avatar_definition, [new_symbols(naming, [spl72_362])])).
fof(f2457, plain, (spl72_358 <=> sP42), introduced(avatar_definition, [new_symbols(naming, [spl72_358])])).
fof(f2438, plain, (spl72_354 <=> sP43), introduced(avatar_definition, [new_symbols(naming, [spl72_354])])).
fof(f2419, plain, (spl72_350 <=> sP44), introduced(avatar_definition, [new_symbols(naming, [spl72_350])])).
fof(f636, plain, (~ (e23 = op2(e23, op2(e23, e23))) | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30), inference(cnf_transformation, [], [f82])).
fof(f82, plain, (((~ (e23 = op2(e23, op2(e23, e23))) & ~ (e22 = op2(e22, op2(e23, e23))) & ~ (e21 = op2(e21, op2(e23, e23))) & ~ (e20 = op2(e20, op2(e23, e23))) & (op2(e23, e23) = op2(e23, e23))) | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45) & ((~ (e23 = op2(op2(e23, e23), e23)) & (e23 = op2(op2(e23, e23), e23)) & ~ (e23 = op2(e23, op2(e23, e23)))) | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30)), inference(definition_folding, [], [f11, e81, e80, e79, e78, e77, e76, e75, e74, e73, e72, e71, e70, e69, e68, e67, e66, e65, e64, e63, e62, e61, e60, e59, e58, e57, e56, e55, e54, e53, e52])).
fof(f52, plain, ((~ (e20 = op2(op2(e20, e20), e20)) & (e20 = op2(op2(e20, e20), e20)) & ~ (e20 = op2(e20, op2(e20, e20)))) | ~ sP30), inference(usedef, [], [e52])).
fof(e52, plain, (sP30 <=> (~ (e20 = op2(op2(e20, e20), e20)) & (e20 = op2(op2(e20, e20), e20)) & ~ (e20 = op2(e20, op2(e20, e20))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP30])])).
fof(f53, plain, ((~ (e20 = op2(op2(e20, e21), e21)) & (e21 = op2(op2(e20, e21), e20)) & ~ (e21 = op2(e20, op2(e20, e21)))) | ~ sP31), inference(usedef, [], [e53])).
fof(e53, plain, (sP31 <=> (~ (e20 = op2(op2(e20, e21), e21)) & (e21 = op2(op2(e20, e21), e20)) & ~ (e21 = op2(e20, op2(e20, e21))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP31])])).
fof(f54, plain, ((~ (e20 = op2(op2(e20, e22), e22)) & (e22 = op2(op2(e20, e22), e20)) & ~ (e22 = op2(e20, op2(e20, e22)))) | ~ sP32), inference(usedef, [], [e54])).
fof(e54, plain, (sP32 <=> (~ (e20 = op2(op2(e20, e22), e22)) & (e22 = op2(op2(e20, e22), e20)) & ~ (e22 = op2(e20, op2(e20, e22))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP32])])).
fof(f55, plain, ((~ (e20 = op2(op2(e20, e23), e23)) & (e23 = op2(op2(e20, e23), e20)) & ~ (e23 = op2(e20, op2(e20, e23)))) | ~ sP33), inference(usedef, [], [e55])).
fof(e55, plain, (sP33 <=> (~ (e20 = op2(op2(e20, e23), e23)) & (e23 = op2(op2(e20, e23), e20)) & ~ (e23 = op2(e20, op2(e20, e23))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP33])])).
fof(f56, plain, ((~ (e21 = op2(op2(e21, e20), e20)) & (e20 = op2(op2(e21, e20), e21)) & ~ (e20 = op2(e21, op2(e21, e20)))) | ~ sP34), inference(usedef, [], [e56])).
fof(e56, plain, (sP34 <=> (~ (e21 = op2(op2(e21, e20), e20)) & (e20 = op2(op2(e21, e20), e21)) & ~ (e20 = op2(e21, op2(e21, e20))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP34])])).
fof(f57, plain, ((~ (e21 = op2(op2(e21, e21), e21)) & (e21 = op2(op2(e21, e21), e21)) & ~ (e21 = op2(e21, op2(e21, e21)))) | ~ sP35), inference(usedef, [], [e57])).
fof(e57, plain, (sP35 <=> (~ (e21 = op2(op2(e21, e21), e21)) & (e21 = op2(op2(e21, e21), e21)) & ~ (e21 = op2(e21, op2(e21, e21))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP35])])).
fof(f58, plain, ((~ (e21 = op2(op2(e21, e22), e22)) & (e22 = op2(op2(e21, e22), e21)) & ~ (e22 = op2(e21, op2(e21, e22)))) | ~ sP36), inference(usedef, [], [e58])).
fof(e58, plain, (sP36 <=> (~ (e21 = op2(op2(e21, e22), e22)) & (e22 = op2(op2(e21, e22), e21)) & ~ (e22 = op2(e21, op2(e21, e22))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP36])])).
fof(f59, plain, ((~ (e21 = op2(op2(e21, e23), e23)) & (e23 = op2(op2(e21, e23), e21)) & ~ (e23 = op2(e21, op2(e21, e23)))) | ~ sP37), inference(usedef, [], [e59])).
fof(e59, plain, (sP37 <=> (~ (e21 = op2(op2(e21, e23), e23)) & (e23 = op2(op2(e21, e23), e21)) & ~ (e23 = op2(e21, op2(e21, e23))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP37])])).
fof(f60, plain, ((~ (e22 = op2(op2(e22, e20), e20)) & (e20 = op2(op2(e22, e20), e22)) & ~ (e20 = op2(e22, op2(e22, e20)))) | ~ sP38), inference(usedef, [], [e60])).
fof(e60, plain, (sP38 <=> (~ (e22 = op2(op2(e22, e20), e20)) & (e20 = op2(op2(e22, e20), e22)) & ~ (e20 = op2(e22, op2(e22, e20))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP38])])).
fof(f61, plain, ((~ (e22 = op2(op2(e22, e21), e21)) & (e21 = op2(op2(e22, e21), e22)) & ~ (e21 = op2(e22, op2(e22, e21)))) | ~ sP39), inference(usedef, [], [e61])).
fof(e61, plain, (sP39 <=> (~ (e22 = op2(op2(e22, e21), e21)) & (e21 = op2(op2(e22, e21), e22)) & ~ (e21 = op2(e22, op2(e22, e21))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP39])])).
fof(f62, plain, ((~ (e22 = op2(op2(e22, e22), e22)) & (e22 = op2(op2(e22, e22), e22)) & ~ (e22 = op2(e22, op2(e22, e22)))) | ~ sP40), inference(usedef, [], [e62])).
fof(e62, plain, (sP40 <=> (~ (e22 = op2(op2(e22, e22), e22)) & (e22 = op2(op2(e22, e22), e22)) & ~ (e22 = op2(e22, op2(e22, e22))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP40])])).
fof(f63, plain, ((~ (e22 = op2(op2(e22, e23), e23)) & (e23 = op2(op2(e22, e23), e22)) & ~ (e23 = op2(e22, op2(e22, e23)))) | ~ sP41), inference(usedef, [], [e63])).
fof(e63, plain, (sP41 <=> (~ (e22 = op2(op2(e22, e23), e23)) & (e23 = op2(op2(e22, e23), e22)) & ~ (e23 = op2(e22, op2(e22, e23))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP41])])).
fof(f64, plain, ((~ (e23 = op2(op2(e23, e20), e20)) & (e20 = op2(op2(e23, e20), e23)) & ~ (e20 = op2(e23, op2(e23, e20)))) | ~ sP42), inference(usedef, [], [e64])).
fof(e64, plain, (sP42 <=> (~ (e23 = op2(op2(e23, e20), e20)) & (e20 = op2(op2(e23, e20), e23)) & ~ (e20 = op2(e23, op2(e23, e20))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP42])])).
fof(f65, plain, ((~ (e23 = op2(op2(e23, e21), e21)) & (e21 = op2(op2(e23, e21), e23)) & ~ (e21 = op2(e23, op2(e23, e21)))) | ~ sP43), inference(usedef, [], [e65])).
fof(e65, plain, (sP43 <=> (~ (e23 = op2(op2(e23, e21), e21)) & (e21 = op2(op2(e23, e21), e23)) & ~ (e21 = op2(e23, op2(e23, e21))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP43])])).
fof(f66, plain, ((~ (e23 = op2(op2(e23, e22), e22)) & (e22 = op2(op2(e23, e22), e23)) & ~ (e22 = op2(e23, op2(e23, e22)))) | ~ sP44), inference(usedef, [], [e66])).
fof(e66, plain, (sP44 <=> (~ (e23 = op2(op2(e23, e22), e22)) & (e22 = op2(op2(e23, e22), e23)) & ~ (e22 = op2(e23, op2(e23, e22))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP44])])).
fof(f67, plain, ((~ (e23 = op2(e23, op2(e20, e20))) & ~ (e22 = op2(e22, op2(e20, e20))) & ~ (e21 = op2(e21, op2(e20, e20))) & ~ (e20 = op2(e20, op2(e20, e20))) & (op2(e20, e20) = op2(e20, e20))) | ~ sP45), inference(usedef, [], [e67])).
fof(e67, plain, (sP45 <=> (~ (e23 = op2(e23, op2(e20, e20))) & ~ (e22 = op2(e22, op2(e20, e20))) & ~ (e21 = op2(e21, op2(e20, e20))) & ~ (e20 = op2(e20, op2(e20, e20))) & (op2(e20, e20) = op2(e20, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP45])])).
fof(f68, plain, ((~ (e23 = op2(e23, op2(e20, e21))) & ~ (e22 = op2(e22, op2(e20, e21))) & ~ (e21 = op2(e21, op2(e20, e21))) & ~ (e20 = op2(e20, op2(e20, e21))) & (op2(e20, e21) = op2(e21, e20))) | ~ sP46), inference(usedef, [], [e68])).
fof(e68, plain, (sP46 <=> (~ (e23 = op2(e23, op2(e20, e21))) & ~ (e22 = op2(e22, op2(e20, e21))) & ~ (e21 = op2(e21, op2(e20, e21))) & ~ (e20 = op2(e20, op2(e20, e21))) & (op2(e20, e21) = op2(e21, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP46])])).
fof(f69, plain, ((~ (e23 = op2(e23, op2(e20, e22))) & ~ (e22 = op2(e22, op2(e20, e22))) & ~ (e21 = op2(e21, op2(e20, e22))) & ~ (e20 = op2(e20, op2(e20, e22))) & (op2(e20, e22) = op2(e22, e20))) | ~ sP47), inference(usedef, [], [e69])).
fof(e69, plain, (sP47 <=> (~ (e23 = op2(e23, op2(e20, e22))) & ~ (e22 = op2(e22, op2(e20, e22))) & ~ (e21 = op2(e21, op2(e20, e22))) & ~ (e20 = op2(e20, op2(e20, e22))) & (op2(e20, e22) = op2(e22, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP47])])).
fof(f70, plain, ((~ (e23 = op2(e23, op2(e20, e23))) & ~ (e22 = op2(e22, op2(e20, e23))) & ~ (e21 = op2(e21, op2(e20, e23))) & ~ (e20 = op2(e20, op2(e20, e23))) & (op2(e20, e23) = op2(e23, e20))) | ~ sP48), inference(usedef, [], [e70])).
fof(e70, plain, (sP48 <=> (~ (e23 = op2(e23, op2(e20, e23))) & ~ (e22 = op2(e22, op2(e20, e23))) & ~ (e21 = op2(e21, op2(e20, e23))) & ~ (e20 = op2(e20, op2(e20, e23))) & (op2(e20, e23) = op2(e23, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP48])])).
fof(f71, plain, ((~ (e23 = op2(e23, op2(e21, e20))) & ~ (e22 = op2(e22, op2(e21, e20))) & ~ (e21 = op2(e21, op2(e21, e20))) & ~ (e20 = op2(e20, op2(e21, e20))) & (op2(e20, e21) = op2(e21, e20))) | ~ sP49), inference(usedef, [], [e71])).
fof(e71, plain, (sP49 <=> (~ (e23 = op2(e23, op2(e21, e20))) & ~ (e22 = op2(e22, op2(e21, e20))) & ~ (e21 = op2(e21, op2(e21, e20))) & ~ (e20 = op2(e20, op2(e21, e20))) & (op2(e20, e21) = op2(e21, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP49])])).
fof(f72, plain, ((~ (e23 = op2(e23, op2(e21, e21))) & ~ (e22 = op2(e22, op2(e21, e21))) & ~ (e21 = op2(e21, op2(e21, e21))) & ~ (e20 = op2(e20, op2(e21, e21))) & (op2(e21, e21) = op2(e21, e21))) | ~ sP50), inference(usedef, [], [e72])).
fof(e72, plain, (sP50 <=> (~ (e23 = op2(e23, op2(e21, e21))) & ~ (e22 = op2(e22, op2(e21, e21))) & ~ (e21 = op2(e21, op2(e21, e21))) & ~ (e20 = op2(e20, op2(e21, e21))) & (op2(e21, e21) = op2(e21, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP50])])).
fof(f73, plain, ((~ (e23 = op2(e23, op2(e21, e22))) & ~ (e22 = op2(e22, op2(e21, e22))) & ~ (e21 = op2(e21, op2(e21, e22))) & ~ (e20 = op2(e20, op2(e21, e22))) & (op2(e21, e22) = op2(e22, e21))) | ~ sP51), inference(usedef, [], [e73])).
fof(e73, plain, (sP51 <=> (~ (e23 = op2(e23, op2(e21, e22))) & ~ (e22 = op2(e22, op2(e21, e22))) & ~ (e21 = op2(e21, op2(e21, e22))) & ~ (e20 = op2(e20, op2(e21, e22))) & (op2(e21, e22) = op2(e22, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP51])])).
fof(f74, plain, ((~ (e23 = op2(e23, op2(e21, e23))) & ~ (e22 = op2(e22, op2(e21, e23))) & ~ (e21 = op2(e21, op2(e21, e23))) & ~ (e20 = op2(e20, op2(e21, e23))) & (op2(e21, e23) = op2(e23, e21))) | ~ sP52), inference(usedef, [], [e74])).
fof(e74, plain, (sP52 <=> (~ (e23 = op2(e23, op2(e21, e23))) & ~ (e22 = op2(e22, op2(e21, e23))) & ~ (e21 = op2(e21, op2(e21, e23))) & ~ (e20 = op2(e20, op2(e21, e23))) & (op2(e21, e23) = op2(e23, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP52])])).
fof(f75, plain, ((~ (e23 = op2(e23, op2(e22, e20))) & ~ (e22 = op2(e22, op2(e22, e20))) & ~ (e21 = op2(e21, op2(e22, e20))) & ~ (e20 = op2(e20, op2(e22, e20))) & (op2(e20, e22) = op2(e22, e20))) | ~ sP53), inference(usedef, [], [e75])).
fof(e75, plain, (sP53 <=> (~ (e23 = op2(e23, op2(e22, e20))) & ~ (e22 = op2(e22, op2(e22, e20))) & ~ (e21 = op2(e21, op2(e22, e20))) & ~ (e20 = op2(e20, op2(e22, e20))) & (op2(e20, e22) = op2(e22, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP53])])).
fof(f76, plain, ((~ (e23 = op2(e23, op2(e22, e21))) & ~ (e22 = op2(e22, op2(e22, e21))) & ~ (e21 = op2(e21, op2(e22, e21))) & ~ (e20 = op2(e20, op2(e22, e21))) & (op2(e21, e22) = op2(e22, e21))) | ~ sP54), inference(usedef, [], [e76])).
fof(e76, plain, (sP54 <=> (~ (e23 = op2(e23, op2(e22, e21))) & ~ (e22 = op2(e22, op2(e22, e21))) & ~ (e21 = op2(e21, op2(e22, e21))) & ~ (e20 = op2(e20, op2(e22, e21))) & (op2(e21, e22) = op2(e22, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP54])])).
fof(f77, plain, ((~ (e23 = op2(e23, op2(e22, e22))) & ~ (e22 = op2(e22, op2(e22, e22))) & ~ (e21 = op2(e21, op2(e22, e22))) & ~ (e20 = op2(e20, op2(e22, e22))) & (op2(e22, e22) = op2(e22, e22))) | ~ sP55), inference(usedef, [], [e77])).
fof(e77, plain, (sP55 <=> (~ (e23 = op2(e23, op2(e22, e22))) & ~ (e22 = op2(e22, op2(e22, e22))) & ~ (e21 = op2(e21, op2(e22, e22))) & ~ (e20 = op2(e20, op2(e22, e22))) & (op2(e22, e22) = op2(e22, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP55])])).
fof(f78, plain, ((~ (e23 = op2(e23, op2(e22, e23))) & ~ (e22 = op2(e22, op2(e22, e23))) & ~ (e21 = op2(e21, op2(e22, e23))) & ~ (e20 = op2(e20, op2(e22, e23))) & (op2(e22, e23) = op2(e23, e22))) | ~ sP56), inference(usedef, [], [e78])).
fof(e78, plain, (sP56 <=> (~ (e23 = op2(e23, op2(e22, e23))) & ~ (e22 = op2(e22, op2(e22, e23))) & ~ (e21 = op2(e21, op2(e22, e23))) & ~ (e20 = op2(e20, op2(e22, e23))) & (op2(e22, e23) = op2(e23, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP56])])).
fof(f79, plain, ((~ (e23 = op2(e23, op2(e23, e20))) & ~ (e22 = op2(e22, op2(e23, e20))) & ~ (e21 = op2(e21, op2(e23, e20))) & ~ (e20 = op2(e20, op2(e23, e20))) & (op2(e20, e23) = op2(e23, e20))) | ~ sP57), inference(usedef, [], [e79])).
fof(e79, plain, (sP57 <=> (~ (e23 = op2(e23, op2(e23, e20))) & ~ (e22 = op2(e22, op2(e23, e20))) & ~ (e21 = op2(e21, op2(e23, e20))) & ~ (e20 = op2(e20, op2(e23, e20))) & (op2(e20, e23) = op2(e23, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP57])])).
fof(f80, plain, ((~ (e23 = op2(e23, op2(e23, e21))) & ~ (e22 = op2(e22, op2(e23, e21))) & ~ (e21 = op2(e21, op2(e23, e21))) & ~ (e20 = op2(e20, op2(e23, e21))) & (op2(e21, e23) = op2(e23, e21))) | ~ sP58), inference(usedef, [], [e80])).
fof(e80, plain, (sP58 <=> (~ (e23 = op2(e23, op2(e23, e21))) & ~ (e22 = op2(e22, op2(e23, e21))) & ~ (e21 = op2(e21, op2(e23, e21))) & ~ (e20 = op2(e20, op2(e23, e21))) & (op2(e21, e23) = op2(e23, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP58])])).
fof(f81, plain, ((~ (e23 = op2(e23, op2(e23, e22))) & ~ (e22 = op2(e22, op2(e23, e22))) & ~ (e21 = op2(e21, op2(e23, e22))) & ~ (e20 = op2(e20, op2(e23, e22))) & (op2(e22, e23) = op2(e23, e22))) | ~ sP59), inference(usedef, [], [e81])).
fof(e81, plain, (sP59 <=> (~ (e23 = op2(e23, op2(e23, e22))) & ~ (e22 = op2(e22, op2(e23, e22))) & ~ (e21 = op2(e21, op2(e23, e22))) & ~ (e20 = op2(e20, op2(e23, e22))) & (op2(e22, e23) = op2(e23, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP59])])).
fof(f11, plain, (((~ (e23 = op2(e23, op2(e23, e23))) & ~ (e22 = op2(e22, op2(e23, e23))) & ~ (e21 = op2(e21, op2(e23, e23))) & ~ (e20 = op2(e20, op2(e23, e23))) & (op2(e23, e23) = op2(e23, e23))) | (~ (e23 = op2(e23, op2(e23, e22))) & ~ (e22 = op2(e22, op2(e23, e22))) & ~ (e21 = op2(e21, op2(e23, e22))) & ~ (e20 = op2(e20, op2(e23, e22))) & (op2(e22, e23) = op2(e23, e22))) | (~ (e23 = op2(e23, op2(e23, e21))) & ~ (e22 = op2(e22, op2(e23, e21))) & ~ (e21 = op2(e21, op2(e23, e21))) & ~ (e20 = op2(e20, op2(e23, e21))) & (op2(e21, e23) = op2(e23, e21))) | (~ (e23 = op2(e23, op2(e23, e20))) & ~ (e22 = op2(e22, op2(e23, e20))) & ~ (e21 = op2(e21, op2(e23, e20))) & ~ (e20 = op2(e20, op2(e23, e20))) & (op2(e20, e23) = op2(e23, e20))) | (~ (e23 = op2(e23, op2(e22, e23))) & ~ (e22 = op2(e22, op2(e22, e23))) & ~ (e21 = op2(e21, op2(e22, e23))) & ~ (e20 = op2(e20, op2(e22, e23))) & (op2(e22, e23) = op2(e23, e22))) | (~ (e23 = op2(e23, op2(e22, e22))) & ~ (e22 = op2(e22, op2(e22, e22))) & ~ (e21 = op2(e21, op2(e22, e22))) & ~ (e20 = op2(e20, op2(e22, e22))) & (op2(e22, e22) = op2(e22, e22))) | (~ (e23 = op2(e23, op2(e22, e21))) & ~ (e22 = op2(e22, op2(e22, e21))) & ~ (e21 = op2(e21, op2(e22, e21))) & ~ (e20 = op2(e20, op2(e22, e21))) & (op2(e21, e22) = op2(e22, e21))) | (~ (e23 = op2(e23, op2(e22, e20))) & ~ (e22 = op2(e22, op2(e22, e20))) & ~ (e21 = op2(e21, op2(e22, e20))) & ~ (e20 = op2(e20, op2(e22, e20))) & (op2(e20, e22) = op2(e22, e20))) | (~ (e23 = op2(e23, op2(e21, e23))) & ~ (e22 = op2(e22, op2(e21, e23))) & ~ (e21 = op2(e21, op2(e21, e23))) & ~ (e20 = op2(e20, op2(e21, e23))) & (op2(e21, e23) = op2(e23, e21))) | (~ (e23 = op2(e23, op2(e21, e22))) & ~ (e22 = op2(e22, op2(e21, e22))) & ~ (e21 = op2(e21, op2(e21, e22))) & ~ (e20 = op2(e20, op2(e21, e22))) & (op2(e21, e22) = op2(e22, e21))) | (~ (e23 = op2(e23, op2(e21, e21))) & ~ (e22 = op2(e22, op2(e21, e21))) & ~ (e21 = op2(e21, op2(e21, e21))) & ~ (e20 = op2(e20, op2(e21, e21))) & (op2(e21, e21) = op2(e21, e21))) | (~ (e23 = op2(e23, op2(e21, e20))) & ~ (e22 = op2(e22, op2(e21, e20))) & ~ (e21 = op2(e21, op2(e21, e20))) & ~ (e20 = op2(e20, op2(e21, e20))) & (op2(e20, e21) = op2(e21, e20))) | (~ (e23 = op2(e23, op2(e20, e23))) & ~ (e22 = op2(e22, op2(e20, e23))) & ~ (e21 = op2(e21, op2(e20, e23))) & ~ (e20 = op2(e20, op2(e20, e23))) & (op2(e20, e23) = op2(e23, e20))) | (~ (e23 = op2(e23, op2(e20, e22))) & ~ (e22 = op2(e22, op2(e20, e22))) & ~ (e21 = op2(e21, op2(e20, e22))) & ~ (e20 = op2(e20, op2(e20, e22))) & (op2(e20, e22) = op2(e22, e20))) | (~ (e23 = op2(e23, op2(e20, e21))) & ~ (e22 = op2(e22, op2(e20, e21))) & ~ (e21 = op2(e21, op2(e20, e21))) & ~ (e20 = op2(e20, op2(e20, e21))) & (op2(e20, e21) = op2(e21, e20))) | (~ (e23 = op2(e23, op2(e20, e20))) & ~ (e22 = op2(e22, op2(e20, e20))) & ~ (e21 = op2(e21, op2(e20, e20))) & ~ (e20 = op2(e20, op2(e20, e20))) & (op2(e20, e20) = op2(e20, e20)))) & ((~ (e23 = op2(op2(e23, e23), e23)) & (e23 = op2(op2(e23, e23), e23)) & ~ (e23 = op2(e23, op2(e23, e23)))) | (~ (e23 = op2(op2(e23, e22), e22)) & (e22 = op2(op2(e23, e22), e23)) & ~ (e22 = op2(e23, op2(e23, e22)))) | (~ (e23 = op2(op2(e23, e21), e21)) & (e21 = op2(op2(e23, e21), e23)) & ~ (e21 = op2(e23, op2(e23, e21)))) | (~ (e23 = op2(op2(e23, e20), e20)) & (e20 = op2(op2(e23, e20), e23)) & ~ (e20 = op2(e23, op2(e23, e20)))) | (~ (e22 = op2(op2(e22, e23), e23)) & (e23 = op2(op2(e22, e23), e22)) & ~ (e23 = op2(e22, op2(e22, e23)))) | (~ (e22 = op2(op2(e22, e22), e22)) & (e22 = op2(op2(e22, e22), e22)) & ~ (e22 = op2(e22, op2(e22, e22)))) | (~ (e22 = op2(op2(e22, e21), e21)) & (e21 = op2(op2(e22, e21), e22)) & ~ (e21 = op2(e22, op2(e22, e21)))) | (~ (e22 = op2(op2(e22, e20), e20)) & (e20 = op2(op2(e22, e20), e22)) & ~ (e20 = op2(e22, op2(e22, e20)))) | (~ (e21 = op2(op2(e21, e23), e23)) & (e23 = op2(op2(e21, e23), e21)) & ~ (e23 = op2(e21, op2(e21, e23)))) | (~ (e21 = op2(op2(e21, e22), e22)) & (e22 = op2(op2(e21, e22), e21)) & ~ (e22 = op2(e21, op2(e21, e22)))) | (~ (e21 = op2(op2(e21, e21), e21)) & (e21 = op2(op2(e21, e21), e21)) & ~ (e21 = op2(e21, op2(e21, e21)))) | (~ (e21 = op2(op2(e21, e20), e20)) & (e20 = op2(op2(e21, e20), e21)) & ~ (e20 = op2(e21, op2(e21, e20)))) | (~ (e20 = op2(op2(e20, e23), e23)) & (e23 = op2(op2(e20, e23), e20)) & ~ (e23 = op2(e20, op2(e20, e23)))) | (~ (e20 = op2(op2(e20, e22), e22)) & (e22 = op2(op2(e20, e22), e20)) & ~ (e22 = op2(e20, op2(e20, e22)))) | (~ (e20 = op2(op2(e20, e21), e21)) & (e21 = op2(op2(e20, e21), e20)) & ~ (e21 = op2(e20, op2(e20, e21)))) | (~ (e20 = op2(op2(e20, e20), e20)) & (e20 = op2(op2(e20, e20), e20)) & ~ (e20 = op2(e20, op2(e20, e20)))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG104+1.p', ax11)).
fof(f2704, plain, (spl72_402 | spl72_398 | spl72_394 | spl72_390 | spl72_386 | spl72_384 | spl72_380 | spl72_376 | spl72_372 | spl72_368 | spl72_366 | spl72_362 | spl72_358 | spl72_354 | spl72_350 | spl72_408), inference(avatar_split_clause, [], [f637, f2700, f2419, f2438, f2457, f2476, f2495, f2506, f2525, f2544, f2563, f2582, f2593, f2612, f2631, f2650, f2669])).
fof(f2700, plain, (spl72_408 <=> (e23 = op2(op2(e23, e23), e23))), introduced(avatar_definition, [new_symbols(naming, [spl72_408])])).
fof(f637, plain, ((e23 = op2(op2(e23, e23), e23)) | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30), inference(cnf_transformation, [], [f82])).
fof(f2703, plain, (spl72_402 | spl72_398 | spl72_394 | spl72_390 | spl72_386 | spl72_384 | spl72_380 | spl72_376 | spl72_372 | spl72_368 | spl72_366 | spl72_362 | spl72_358 | spl72_354 | spl72_350 | ~ spl72_408), inference(avatar_split_clause, [], [f638, f2700, f2419, f2438, f2457, f2476, f2495, f2506, f2525, f2544, f2563, f2582, f2593, f2612, f2631, f2650, f2669])).
fof(f638, plain, (~ (e23 = op2(op2(e23, e23), e23)) | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30), inference(cnf_transformation, [], [f82])).
fof(f2688, plain, (spl72_345 | spl72_340 | spl72_335 | spl72_330 | spl72_324 | spl72_319 | spl72_314 | spl72_309 | spl72_303 | spl72_297 | spl72_292 | spl72_287 | spl72_281 | spl72_275 | spl72_269 | ~ spl72_405), inference(avatar_split_clause, [], [f642, f2685, f2023, f2052, f2081, f2110, f2135, f2159, f2188, f2217, f2242, f2267, f2291, f2320, f2345, f2370, f2395])).
fof(f2395, plain, (spl72_345 <=> sP45), introduced(avatar_definition, [new_symbols(naming, [spl72_345])])).
fof(f2370, plain, (spl72_340 <=> sP46), introduced(avatar_definition, [new_symbols(naming, [spl72_340])])).
fof(f2345, plain, (spl72_335 <=> sP47), introduced(avatar_definition, [new_symbols(naming, [spl72_335])])).
fof(f2320, plain, (spl72_330 <=> sP48), introduced(avatar_definition, [new_symbols(naming, [spl72_330])])).
fof(f2291, plain, (spl72_324 <=> sP49), introduced(avatar_definition, [new_symbols(naming, [spl72_324])])).
fof(f2267, plain, (spl72_319 <=> sP50), introduced(avatar_definition, [new_symbols(naming, [spl72_319])])).
fof(f2242, plain, (spl72_314 <=> sP51), introduced(avatar_definition, [new_symbols(naming, [spl72_314])])).
fof(f2217, plain, (spl72_309 <=> sP52), introduced(avatar_definition, [new_symbols(naming, [spl72_309])])).
fof(f2188, plain, (spl72_303 <=> sP53), introduced(avatar_definition, [new_symbols(naming, [spl72_303])])).
fof(f2159, plain, (spl72_297 <=> sP54), introduced(avatar_definition, [new_symbols(naming, [spl72_297])])).
fof(f2135, plain, (spl72_292 <=> sP55), introduced(avatar_definition, [new_symbols(naming, [spl72_292])])).
fof(f2110, plain, (spl72_287 <=> sP56), introduced(avatar_definition, [new_symbols(naming, [spl72_287])])).
fof(f2081, plain, (spl72_281 <=> sP57), introduced(avatar_definition, [new_symbols(naming, [spl72_281])])).
fof(f2052, plain, (spl72_275 <=> sP58), introduced(avatar_definition, [new_symbols(naming, [spl72_275])])).
fof(f2023, plain, (spl72_269 <=> sP59), introduced(avatar_definition, [new_symbols(naming, [spl72_269])])).
fof(f642, plain, (~ (e22 = op2(e22, op2(e23, e23))) | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45), inference(cnf_transformation, [], [f82])).
fof(f2683, plain, (spl72_345 | spl72_340 | spl72_335 | spl72_330 | spl72_324 | spl72_319 | spl72_314 | spl72_309 | spl72_303 | spl72_297 | spl72_292 | spl72_287 | spl72_281 | spl72_275 | spl72_269 | ~ spl72_404), inference(avatar_split_clause, [], [f643, f2680, f2023, f2052, f2081, f2110, f2135, f2159, f2188, f2217, f2242, f2267, f2291, f2320, f2345, f2370, f2395])).
fof(f643, plain, (~ (e23 = op2(e23, op2(e23, e23))) | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45), inference(cnf_transformation, [], [f82])).
fof(f2677, plain, (~ spl72_402 | spl72_403), inference(avatar_split_clause, [], [f634, f2673, f2669])).
fof(f2673, plain, (spl72_403 <=> (e20 = op2(op2(e20, e20), e20))), introduced(avatar_definition, [new_symbols(naming, [spl72_403])])).
fof(f634, plain, ((e20 = op2(op2(e20, e20), e20)) | ~ sP30), inference(cnf_transformation, [], [f155])).
fof(f155, plain, ((~ (e20 = op2(op2(e20, e20), e20)) & (e20 = op2(op2(e20, e20), e20)) & ~ (e20 = op2(e20, op2(e20, e20)))) | ~ sP30), inference(nnf_transformation, [], [f52])).
fof(f2676, plain, (~ spl72_402 | ~ spl72_403), inference(avatar_split_clause, [], [f635, f2673, f2669])).
fof(f635, plain, (~ (e20 = op2(op2(e20, e20), e20)) | ~ sP30), inference(cnf_transformation, [], [f155])).
fof(f2662, plain, (~ spl72_398 | spl72_400), inference(avatar_split_clause, [], [f631, f2659, f2650])).
fof(f631, plain, ((e21 = op2(op2(e20, e21), e20)) | ~ sP31), inference(cnf_transformation, [], [f154])).
fof(f154, plain, ((~ (e20 = op2(op2(e20, e21), e21)) & (e21 = op2(op2(e20, e21), e20)) & ~ (e21 = op2(e20, op2(e20, e21)))) | ~ sP31), inference(nnf_transformation, [], [f53])).
fof(f2657, plain, (~ spl72_398 | ~ spl72_399), inference(avatar_split_clause, [], [f632, f2654, f2650])).
fof(f632, plain, (~ (e20 = op2(op2(e20, e21), e21)) | ~ sP31), inference(cnf_transformation, [], [f154])).
fof(f2643, plain, (~ spl72_394 | spl72_396), inference(avatar_split_clause, [], [f628, f2640, f2631])).
fof(f628, plain, ((e22 = op2(op2(e20, e22), e20)) | ~ sP32), inference(cnf_transformation, [], [f153])).
fof(f153, plain, ((~ (e20 = op2(op2(e20, e22), e22)) & (e22 = op2(op2(e20, e22), e20)) & ~ (e22 = op2(e20, op2(e20, e22)))) | ~ sP32), inference(nnf_transformation, [], [f54])).
fof(f2619, plain, (~ spl72_390 | ~ spl72_391), inference(avatar_split_clause, [], [f626, f2616, f2612])).
fof(f626, plain, (~ (e20 = op2(op2(e20, e23), e23)) | ~ sP33), inference(cnf_transformation, [], [f152])).
fof(f152, plain, ((~ (e20 = op2(op2(e20, e23), e23)) & (e23 = op2(op2(e20, e23), e20)) & ~ (e23 = op2(e20, op2(e20, e23)))) | ~ sP33), inference(nnf_transformation, [], [f55])).
fof(f2610, plain, (~ spl72_386 | ~ spl72_389), inference(avatar_split_clause, [], [f621, f2607, f2593])).
fof(f621, plain, (~ (e20 = op2(e21, op2(e21, e20))) | ~ sP34), inference(cnf_transformation, [], [f151])).
fof(f151, plain, ((~ (e21 = op2(op2(e21, e20), e20)) & (e20 = op2(op2(e21, e20), e21)) & ~ (e20 = op2(e21, op2(e21, e20)))) | ~ sP34), inference(nnf_transformation, [], [f56])).
fof(f2605, plain, (~ spl72_386 | spl72_388), inference(avatar_split_clause, [], [f622, f2602, f2593])).
fof(f622, plain, ((e20 = op2(op2(e21, e20), e21)) | ~ sP34), inference(cnf_transformation, [], [f151])).
fof(f2600, plain, (~ spl72_386 | ~ spl72_387), inference(avatar_split_clause, [], [f623, f2597, f2593])).
fof(f623, plain, (~ (e21 = op2(op2(e21, e20), e20)) | ~ sP34), inference(cnf_transformation, [], [f151])).
fof(f2590, plain, (~ spl72_384 | spl72_385), inference(avatar_split_clause, [], [f619, f2586, f2582])).
fof(f2586, plain, (spl72_385 <=> (e21 = op2(op2(e21, e21), e21))), introduced(avatar_definition, [new_symbols(naming, [spl72_385])])).
fof(f619, plain, ((e21 = op2(op2(e21, e21), e21)) | ~ sP35), inference(cnf_transformation, [], [f150])).
fof(f150, plain, ((~ (e21 = op2(op2(e21, e21), e21)) & (e21 = op2(op2(e21, e21), e21)) & ~ (e21 = op2(e21, op2(e21, e21)))) | ~ sP35), inference(nnf_transformation, [], [f57])).
fof(f2589, plain, (~ spl72_384 | ~ spl72_385), inference(avatar_split_clause, [], [f620, f2586, f2582])).
fof(f620, plain, (~ (e21 = op2(op2(e21, e21), e21)) | ~ sP35), inference(cnf_transformation, [], [f150])).
fof(f2580, plain, (~ spl72_380 | ~ spl72_383), inference(avatar_split_clause, [], [f615, f2577, f2563])).
fof(f615, plain, (~ (e22 = op2(e21, op2(e21, e22))) | ~ sP36), inference(cnf_transformation, [], [f149])).
fof(f149, plain, ((~ (e21 = op2(op2(e21, e22), e22)) & (e22 = op2(op2(e21, e22), e21)) & ~ (e22 = op2(e21, op2(e21, e22)))) | ~ sP36), inference(nnf_transformation, [], [f58])).
fof(f2575, plain, (~ spl72_380 | spl72_382), inference(avatar_split_clause, [], [f616, f2572, f2563])).
fof(f616, plain, ((e22 = op2(op2(e21, e22), e21)) | ~ sP36), inference(cnf_transformation, [], [f149])).
fof(f2570, plain, (~ spl72_380 | ~ spl72_381), inference(avatar_split_clause, [], [f617, f2567, f2563])).
fof(f617, plain, (~ (e21 = op2(op2(e21, e22), e22)) | ~ sP36), inference(cnf_transformation, [], [f149])).
fof(f2561, plain, (~ spl72_376 | ~ spl72_379), inference(avatar_split_clause, [], [f612, f2558, f2544])).
fof(f612, plain, (~ (e23 = op2(e21, op2(e21, e23))) | ~ sP37), inference(cnf_transformation, [], [f148])).
fof(f148, plain, ((~ (e21 = op2(op2(e21, e23), e23)) & (e23 = op2(op2(e21, e23), e21)) & ~ (e23 = op2(e21, op2(e21, e23)))) | ~ sP37), inference(nnf_transformation, [], [f59])).
fof(f2556, plain, (~ spl72_376 | spl72_378), inference(avatar_split_clause, [], [f613, f2553, f2544])).
fof(f613, plain, ((e23 = op2(op2(e21, e23), e21)) | ~ sP37), inference(cnf_transformation, [], [f148])).
fof(f2551, plain, (~ spl72_376 | ~ spl72_377), inference(avatar_split_clause, [], [f614, f2548, f2544])).
fof(f614, plain, (~ (e21 = op2(op2(e21, e23), e23)) | ~ sP37), inference(cnf_transformation, [], [f148])).
fof(f2542, plain, (~ spl72_372 | ~ spl72_375), inference(avatar_split_clause, [], [f609, f2539, f2525])).
fof(f609, plain, (~ (e20 = op2(e22, op2(e22, e20))) | ~ sP38), inference(cnf_transformation, [], [f147])).
fof(f147, plain, ((~ (e22 = op2(op2(e22, e20), e20)) & (e20 = op2(op2(e22, e20), e22)) & ~ (e20 = op2(e22, op2(e22, e20)))) | ~ sP38), inference(nnf_transformation, [], [f60])).
fof(f2537, plain, (~ spl72_372 | spl72_374), inference(avatar_split_clause, [], [f610, f2534, f2525])).
fof(f610, plain, ((e20 = op2(op2(e22, e20), e22)) | ~ sP38), inference(cnf_transformation, [], [f147])).
fof(f2532, plain, (~ spl72_372 | ~ spl72_373), inference(avatar_split_clause, [], [f611, f2529, f2525])).
fof(f611, plain, (~ (e22 = op2(op2(e22, e20), e20)) | ~ sP38), inference(cnf_transformation, [], [f147])).
fof(f2523, plain, (~ spl72_368 | ~ spl72_371), inference(avatar_split_clause, [], [f606, f2520, f2506])).
fof(f606, plain, (~ (e21 = op2(e22, op2(e22, e21))) | ~ sP39), inference(cnf_transformation, [], [f146])).
fof(f146, plain, ((~ (e22 = op2(op2(e22, e21), e21)) & (e21 = op2(op2(e22, e21), e22)) & ~ (e21 = op2(e22, op2(e22, e21)))) | ~ sP39), inference(nnf_transformation, [], [f61])).
fof(f2518, plain, (~ spl72_368 | spl72_370), inference(avatar_split_clause, [], [f607, f2515, f2506])).
fof(f607, plain, ((e21 = op2(op2(e22, e21), e22)) | ~ sP39), inference(cnf_transformation, [], [f146])).
fof(f2513, plain, (~ spl72_368 | ~ spl72_369), inference(avatar_split_clause, [], [f608, f2510, f2506])).
fof(f608, plain, (~ (e22 = op2(op2(e22, e21), e21)) | ~ sP39), inference(cnf_transformation, [], [f146])).
fof(f2503, plain, (~ spl72_366 | spl72_367), inference(avatar_split_clause, [], [f604, f2499, f2495])).
fof(f2499, plain, (spl72_367 <=> (e22 = op2(op2(e22, e22), e22))), introduced(avatar_definition, [new_symbols(naming, [spl72_367])])).
fof(f604, plain, ((e22 = op2(op2(e22, e22), e22)) | ~ sP40), inference(cnf_transformation, [], [f145])).
fof(f145, plain, ((~ (e22 = op2(op2(e22, e22), e22)) & (e22 = op2(op2(e22, e22), e22)) & ~ (e22 = op2(e22, op2(e22, e22)))) | ~ sP40), inference(nnf_transformation, [], [f62])).
fof(f2502, plain, (~ spl72_366 | ~ spl72_367), inference(avatar_split_clause, [], [f605, f2499, f2495])).
fof(f605, plain, (~ (e22 = op2(op2(e22, e22), e22)) | ~ sP40), inference(cnf_transformation, [], [f145])).
fof(f2493, plain, (~ spl72_362 | ~ spl72_365), inference(avatar_split_clause, [], [f600, f2490, f2476])).
fof(f600, plain, (~ (e23 = op2(e22, op2(e22, e23))) | ~ sP41), inference(cnf_transformation, [], [f144])).
fof(f144, plain, ((~ (e22 = op2(op2(e22, e23), e23)) & (e23 = op2(op2(e22, e23), e22)) & ~ (e23 = op2(e22, op2(e22, e23)))) | ~ sP41), inference(nnf_transformation, [], [f63])).
fof(f2488, plain, (~ spl72_362 | spl72_364), inference(avatar_split_clause, [], [f601, f2485, f2476])).
fof(f601, plain, ((e23 = op2(op2(e22, e23), e22)) | ~ sP41), inference(cnf_transformation, [], [f144])).
fof(f2483, plain, (~ spl72_362 | ~ spl72_363), inference(avatar_split_clause, [], [f602, f2480, f2476])).
fof(f602, plain, (~ (e22 = op2(op2(e22, e23), e23)) | ~ sP41), inference(cnf_transformation, [], [f144])).
fof(f2474, plain, (~ spl72_358 | ~ spl72_361), inference(avatar_split_clause, [], [f597, f2471, f2457])).
fof(f597, plain, (~ (e20 = op2(e23, op2(e23, e20))) | ~ sP42), inference(cnf_transformation, [], [f143])).
fof(f143, plain, ((~ (e23 = op2(op2(e23, e20), e20)) & (e20 = op2(op2(e23, e20), e23)) & ~ (e20 = op2(e23, op2(e23, e20)))) | ~ sP42), inference(nnf_transformation, [], [f64])).
fof(f2469, plain, (~ spl72_358 | spl72_360), inference(avatar_split_clause, [], [f598, f2466, f2457])).
fof(f598, plain, ((e20 = op2(op2(e23, e20), e23)) | ~ sP42), inference(cnf_transformation, [], [f143])).
fof(f2464, plain, (~ spl72_358 | ~ spl72_359), inference(avatar_split_clause, [], [f599, f2461, f2457])).
fof(f599, plain, (~ (e23 = op2(op2(e23, e20), e20)) | ~ sP42), inference(cnf_transformation, [], [f143])).
fof(f2455, plain, (~ spl72_354 | ~ spl72_357), inference(avatar_split_clause, [], [f594, f2452, f2438])).
fof(f594, plain, (~ (e21 = op2(e23, op2(e23, e21))) | ~ sP43), inference(cnf_transformation, [], [f142])).
fof(f142, plain, ((~ (e23 = op2(op2(e23, e21), e21)) & (e21 = op2(op2(e23, e21), e23)) & ~ (e21 = op2(e23, op2(e23, e21)))) | ~ sP43), inference(nnf_transformation, [], [f65])).
fof(f2450, plain, (~ spl72_354 | spl72_356), inference(avatar_split_clause, [], [f595, f2447, f2438])).
fof(f595, plain, ((e21 = op2(op2(e23, e21), e23)) | ~ sP43), inference(cnf_transformation, [], [f142])).
fof(f2445, plain, (~ spl72_354 | ~ spl72_355), inference(avatar_split_clause, [], [f596, f2442, f2438])).
fof(f596, plain, (~ (e23 = op2(op2(e23, e21), e21)) | ~ sP43), inference(cnf_transformation, [], [f142])).
fof(f2436, plain, (~ spl72_350 | ~ spl72_353), inference(avatar_split_clause, [], [f591, f2433, f2419])).
fof(f591, plain, (~ (e22 = op2(e23, op2(e23, e22))) | ~ sP44), inference(cnf_transformation, [], [f141])).
fof(f141, plain, ((~ (e23 = op2(op2(e23, e22), e22)) & (e22 = op2(op2(e23, e22), e23)) & ~ (e22 = op2(e23, op2(e23, e22)))) | ~ sP44), inference(nnf_transformation, [], [f66])).
fof(f2431, plain, (~ spl72_350 | spl72_352), inference(avatar_split_clause, [], [f592, f2428, f2419])).
fof(f592, plain, ((e22 = op2(op2(e23, e22), e23)) | ~ sP44), inference(cnf_transformation, [], [f141])).
fof(f2426, plain, (~ spl72_350 | ~ spl72_351), inference(avatar_split_clause, [], [f593, f2423, f2419])).
fof(f593, plain, (~ (e23 = op2(op2(e23, e22), e22)) | ~ sP44), inference(cnf_transformation, [], [f141])).
fof(f2407, plain, (~ spl72_345 | ~ spl72_347), inference(avatar_split_clause, [], [f589, f2404, f2395])).
fof(f589, plain, (~ (e22 = op2(e22, op2(e20, e20))) | ~ sP45), inference(cnf_transformation, [], [f140])).
fof(f140, plain, ((~ (e23 = op2(e23, op2(e20, e20))) & ~ (e22 = op2(e22, op2(e20, e20))) & ~ (e21 = op2(e21, op2(e20, e20))) & ~ (e20 = op2(e20, op2(e20, e20))) & (op2(e20, e20) = op2(e20, e20))) | ~ sP45), inference(nnf_transformation, [], [f67])).
fof(f2393, plain, (~ spl72_340 | spl72_329), inference(avatar_split_clause, [], [f581, f2315, f2370])).
fof(f581, plain, ((op2(e20, e21) = op2(e21, e20)) | ~ sP46), inference(cnf_transformation, [], [f139])).
fof(f139, plain, ((~ (e23 = op2(e23, op2(e20, e21))) & ~ (e22 = op2(e22, op2(e20, e21))) & ~ (e21 = op2(e21, op2(e20, e21))) & ~ (e20 = op2(e20, op2(e20, e21))) & (op2(e20, e21) = op2(e21, e20))) | ~ sP46), inference(nnf_transformation, [], [f68])).
fof(f2368, plain, (~ spl72_335 | spl72_308), inference(avatar_split_clause, [], [f576, f2212, f2345])).
fof(f576, plain, ((op2(e20, e22) = op2(e22, e20)) | ~ sP47), inference(cnf_transformation, [], [f138])).
fof(f138, plain, ((~ (e23 = op2(e23, op2(e20, e22))) & ~ (e22 = op2(e22, op2(e20, e22))) & ~ (e21 = op2(e21, op2(e20, e22))) & ~ (e20 = op2(e20, op2(e20, e22))) & (op2(e20, e22) = op2(e22, e20))) | ~ sP47), inference(nnf_transformation, [], [f69])).
fof(f2343, plain, (~ spl72_330 | spl72_286), inference(avatar_split_clause, [], [f571, f2105, f2320])).
fof(f571, plain, ((op2(e20, e23) = op2(e23, e20)) | ~ sP48), inference(cnf_transformation, [], [f137])).
fof(f137, plain, ((~ (e23 = op2(e23, op2(e20, e23))) & ~ (e22 = op2(e22, op2(e20, e23))) & ~ (e21 = op2(e21, op2(e20, e23))) & ~ (e20 = op2(e20, op2(e20, e23))) & (op2(e20, e23) = op2(e23, e20))) | ~ sP48), inference(nnf_transformation, [], [f70])).
fof(f2337, plain, (~ spl72_330 | ~ spl72_333), inference(avatar_split_clause, [], [f573, f2334, f2320])).
fof(f573, plain, (~ (e21 = op2(e21, op2(e20, e23))) | ~ sP48), inference(cnf_transformation, [], [f137])).
fof(f2318, plain, (~ spl72_324 | spl72_329), inference(avatar_split_clause, [], [f566, f2315, f2291])).
fof(f566, plain, ((op2(e20, e21) = op2(e21, e20)) | ~ sP49), inference(cnf_transformation, [], [f136])).
fof(f136, plain, ((~ (e23 = op2(e23, op2(e21, e20))) & ~ (e22 = op2(e22, op2(e21, e20))) & ~ (e21 = op2(e21, op2(e21, e20))) & ~ (e20 = op2(e20, op2(e21, e20))) & (op2(e20, e21) = op2(e21, e20))) | ~ sP49), inference(nnf_transformation, [], [f71])).
fof(f2284, plain, (~ spl72_319 | ~ spl72_322), inference(avatar_split_clause, [], [f563, f2281, f2267])).
fof(f563, plain, (~ (e21 = op2(e21, op2(e21, e21))) | ~ sP50), inference(cnf_transformation, [], [f135])).
fof(f135, plain, ((~ (e23 = op2(e23, op2(e21, e21))) & ~ (e22 = op2(e22, op2(e21, e21))) & ~ (e21 = op2(e21, op2(e21, e21))) & ~ (e20 = op2(e20, op2(e21, e21))) & (op2(e21, e21) = op2(e21, e21))) | ~ sP50), inference(nnf_transformation, [], [f72])).
fof(f2279, plain, (~ spl72_319 | ~ spl72_321), inference(avatar_split_clause, [], [f564, f2276, f2267])).
fof(f564, plain, (~ (e22 = op2(e22, op2(e21, e21))) | ~ sP50), inference(cnf_transformation, [], [f135])).
fof(f2265, plain, (~ spl72_314 | spl72_302), inference(avatar_split_clause, [], [f556, f2183, f2242])).
fof(f556, plain, ((op2(e21, e22) = op2(e22, e21)) | ~ sP51), inference(cnf_transformation, [], [f134])).
fof(f134, plain, ((~ (e23 = op2(e23, op2(e21, e22))) & ~ (e22 = op2(e22, op2(e21, e22))) & ~ (e21 = op2(e21, op2(e21, e22))) & ~ (e20 = op2(e20, op2(e21, e22))) & (op2(e21, e22) = op2(e22, e21))) | ~ sP51), inference(nnf_transformation, [], [f73])).
fof(f2259, plain, (~ spl72_314 | ~ spl72_317), inference(avatar_split_clause, [], [f558, f2256, f2242])).
fof(f558, plain, (~ (e21 = op2(e21, op2(e21, e22))) | ~ sP51), inference(cnf_transformation, [], [f134])).
fof(f2240, plain, (~ spl72_309 | spl72_280), inference(avatar_split_clause, [], [f551, f2076, f2217])).
fof(f551, plain, ((op2(e21, e23) = op2(e23, e21)) | ~ sP52), inference(cnf_transformation, [], [f133])).
fof(f133, plain, ((~ (e23 = op2(e23, op2(e21, e23))) & ~ (e22 = op2(e22, op2(e21, e23))) & ~ (e21 = op2(e21, op2(e21, e23))) & ~ (e20 = op2(e20, op2(e21, e23))) & (op2(e21, e23) = op2(e23, e21))) | ~ sP52), inference(nnf_transformation, [], [f74])).
fof(f2229, plain, (~ spl72_309 | ~ spl72_311), inference(avatar_split_clause, [], [f554, f2226, f2217])).
fof(f554, plain, (~ (e22 = op2(e22, op2(e21, e23))) | ~ sP52), inference(cnf_transformation, [], [f133])).
fof(f2215, plain, (~ spl72_303 | spl72_308), inference(avatar_split_clause, [], [f546, f2212, f2188])).
fof(f546, plain, ((op2(e20, e22) = op2(e22, e20)) | ~ sP53), inference(cnf_transformation, [], [f132])).
fof(f132, plain, ((~ (e23 = op2(e23, op2(e22, e20))) & ~ (e22 = op2(e22, op2(e22, e20))) & ~ (e21 = op2(e21, op2(e22, e20))) & ~ (e20 = op2(e20, op2(e22, e20))) & (op2(e20, e22) = op2(e22, e20))) | ~ sP53), inference(nnf_transformation, [], [f75])).
fof(f2210, plain, (~ spl72_303 | ~ spl72_307), inference(avatar_split_clause, [], [f547, f2207, f2188])).
fof(f547, plain, (~ (e20 = op2(e20, op2(e22, e20))) | ~ sP53), inference(cnf_transformation, [], [f132])).
fof(f2186, plain, (~ spl72_297 | spl72_302), inference(avatar_split_clause, [], [f541, f2183, f2159])).
fof(f541, plain, ((op2(e21, e22) = op2(e22, e21)) | ~ sP54), inference(cnf_transformation, [], [f131])).
fof(f131, plain, ((~ (e23 = op2(e23, op2(e22, e21))) & ~ (e22 = op2(e22, op2(e22, e21))) & ~ (e21 = op2(e21, op2(e22, e21))) & ~ (e20 = op2(e20, op2(e22, e21))) & (op2(e21, e22) = op2(e22, e21))) | ~ sP54), inference(nnf_transformation, [], [f76])).
fof(f2176, plain, (~ spl72_297 | ~ spl72_300), inference(avatar_split_clause, [], [f543, f2173, f2159])).
fof(f543, plain, (~ (e21 = op2(e21, op2(e22, e21))) | ~ sP54), inference(cnf_transformation, [], [f131])).
fof(f2147, plain, (~ spl72_292 | ~ spl72_294), inference(avatar_split_clause, [], [f539, f2144, f2135])).
fof(f539, plain, (~ (e22 = op2(e22, op2(e22, e22))) | ~ sP55), inference(cnf_transformation, [], [f130])).
fof(f130, plain, ((~ (e23 = op2(e23, op2(e22, e22))) & ~ (e22 = op2(e22, op2(e22, e22))) & ~ (e21 = op2(e21, op2(e22, e22))) & ~ (e20 = op2(e20, op2(e22, e22))) & (op2(e22, e22) = op2(e22, e22))) | ~ sP55), inference(nnf_transformation, [], [f77])).
fof(f2133, plain, (~ spl72_287 | spl72_274), inference(avatar_split_clause, [], [f531, f2047, f2110])).
fof(f531, plain, ((op2(e22, e23) = op2(e23, e22)) | ~ sP56), inference(cnf_transformation, [], [f129])).
fof(f129, plain, ((~ (e23 = op2(e23, op2(e22, e23))) & ~ (e22 = op2(e22, op2(e22, e23))) & ~ (e21 = op2(e21, op2(e22, e23))) & ~ (e20 = op2(e20, op2(e22, e23))) & (op2(e22, e23) = op2(e23, e22))) | ~ sP56), inference(nnf_transformation, [], [f78])).
fof(f2108, plain, (~ spl72_281 | spl72_286), inference(avatar_split_clause, [], [f526, f2105, f2081])).
fof(f526, plain, ((op2(e20, e23) = op2(e23, e20)) | ~ sP57), inference(cnf_transformation, [], [f128])).
fof(f128, plain, ((~ (e23 = op2(e23, op2(e23, e20))) & ~ (e22 = op2(e22, op2(e23, e20))) & ~ (e21 = op2(e21, op2(e23, e20))) & ~ (e20 = op2(e20, op2(e23, e20))) & (op2(e20, e23) = op2(e23, e20))) | ~ sP57), inference(nnf_transformation, [], [f79])).
fof(f2098, plain, (~ spl72_281 | ~ spl72_284), inference(avatar_split_clause, [], [f528, f2095, f2081])).
fof(f528, plain, (~ (e21 = op2(e21, op2(e23, e20))) | ~ sP57), inference(cnf_transformation, [], [f128])).
fof(f2079, plain, (~ spl72_275 | spl72_280), inference(avatar_split_clause, [], [f521, f2076, f2052])).
fof(f521, plain, ((op2(e21, e23) = op2(e23, e21)) | ~ sP58), inference(cnf_transformation, [], [f127])).
fof(f127, plain, ((~ (e23 = op2(e23, op2(e23, e21))) & ~ (e22 = op2(e22, op2(e23, e21))) & ~ (e21 = op2(e21, op2(e23, e21))) & ~ (e20 = op2(e20, op2(e23, e21))) & (op2(e21, e23) = op2(e23, e21))) | ~ sP58), inference(nnf_transformation, [], [f80])).
fof(f2064, plain, (~ spl72_275 | ~ spl72_277), inference(avatar_split_clause, [], [f524, f2061, f2052])).
fof(f524, plain, (~ (e22 = op2(e22, op2(e23, e21))) | ~ sP58), inference(cnf_transformation, [], [f127])).
fof(f2050, plain, (~ spl72_269 | spl72_274), inference(avatar_split_clause, [], [f516, f2047, f2023])).
fof(f516, plain, ((op2(e22, e23) = op2(e23, e22)) | ~ sP59), inference(cnf_transformation, [], [f126])).
fof(f126, plain, ((~ (e23 = op2(e23, op2(e23, e22))) & ~ (e22 = op2(e22, op2(e23, e22))) & ~ (e21 = op2(e21, op2(e23, e22))) & ~ (e20 = op2(e20, op2(e23, e22))) & (op2(e22, e23) = op2(e23, e22))) | ~ sP59), inference(nnf_transformation, [], [f81])).
fof(f2021, plain, (spl72_262 | spl72_258 | spl72_254 | spl72_250 | spl72_246 | spl72_244 | spl72_240 | spl72_236 | spl72_232 | spl72_228 | spl72_226 | spl72_222 | spl72_218 | spl72_214 | spl72_210 | ~ spl72_264), inference(avatar_split_clause, [], [f508, f1996, f1735, f1754, f1773, f1792, f1811, f1822, f1841, f1860, f1879, f1898, f1909, f1928, f1947, f1966, f1985])).
fof(f1985, plain, (spl72_262 <=> sP0), introduced(avatar_definition, [new_symbols(naming, [spl72_262])])).
fof(f1966, plain, (spl72_258 <=> sP1), introduced(avatar_definition, [new_symbols(naming, [spl72_258])])).
fof(f1947, plain, (spl72_254 <=> sP2), introduced(avatar_definition, [new_symbols(naming, [spl72_254])])).
fof(f1928, plain, (spl72_250 <=> sP3), introduced(avatar_definition, [new_symbols(naming, [spl72_250])])).
fof(f1909, plain, (spl72_246 <=> sP4), introduced(avatar_definition, [new_symbols(naming, [spl72_246])])).
fof(f1898, plain, (spl72_244 <=> sP5), introduced(avatar_definition, [new_symbols(naming, [spl72_244])])).
fof(f1879, plain, (spl72_240 <=> sP6), introduced(avatar_definition, [new_symbols(naming, [spl72_240])])).
fof(f1860, plain, (spl72_236 <=> sP7), introduced(avatar_definition, [new_symbols(naming, [spl72_236])])).
fof(f1841, plain, (spl72_232 <=> sP8), introduced(avatar_definition, [new_symbols(naming, [spl72_232])])).
fof(f1822, plain, (spl72_228 <=> sP9), introduced(avatar_definition, [new_symbols(naming, [spl72_228])])).
fof(f1811, plain, (spl72_226 <=> sP10), introduced(avatar_definition, [new_symbols(naming, [spl72_226])])).
fof(f1792, plain, (spl72_222 <=> sP11), introduced(avatar_definition, [new_symbols(naming, [spl72_222])])).
fof(f1773, plain, (spl72_218 <=> sP12), introduced(avatar_definition, [new_symbols(naming, [spl72_218])])).
fof(f1754, plain, (spl72_214 <=> sP13), introduced(avatar_definition, [new_symbols(naming, [spl72_214])])).
fof(f1735, plain, (spl72_210 <=> sP14), introduced(avatar_definition, [new_symbols(naming, [spl72_210])])).
fof(f508, plain, (~ (e13 = op1(e13, op1(e13, e13))) | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f51])).
fof(f51, plain, (((~ (e13 = op1(e13, op1(e13, e13))) & ~ (e12 = op1(e12, op1(e13, e13))) & ~ (e11 = op1(e11, op1(e13, e13))) & ~ (e10 = op1(e10, op1(e13, e13))) & (op1(e13, e13) = op1(e13, e13))) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15) & ((~ (e13 = op1(op1(e13, e13), e13)) & (e13 = op1(op1(e13, e13), e13)) & ~ (e13 = op1(e13, op1(e13, e13)))) | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0)), inference(definition_folding, [], [f10, e50, e49, e48, e47, e46, e45, e44, e43, e42, e41, e40, e39, e38, e37, e36, e35, e34, e33, e32, e31, e30, e29, e28, e27, e26, e25, e24, e23, e22, e21])).
fof(f21, plain, ((~ (e10 = op1(op1(e10, e10), e10)) & (e10 = op1(op1(e10, e10), e10)) & ~ (e10 = op1(e10, op1(e10, e10)))) | ~ sP0), inference(usedef, [], [e21])).
fof(e21, plain, (sP0 <=> (~ (e10 = op1(op1(e10, e10), e10)) & (e10 = op1(op1(e10, e10), e10)) & ~ (e10 = op1(e10, op1(e10, e10))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f22, plain, ((~ (e10 = op1(op1(e10, e11), e11)) & (e11 = op1(op1(e10, e11), e10)) & ~ (e11 = op1(e10, op1(e10, e11)))) | ~ sP1), inference(usedef, [], [e22])).
fof(e22, plain, (sP1 <=> (~ (e10 = op1(op1(e10, e11), e11)) & (e11 = op1(op1(e10, e11), e10)) & ~ (e11 = op1(e10, op1(e10, e11))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f23, plain, ((~ (e10 = op1(op1(e10, e12), e12)) & (e12 = op1(op1(e10, e12), e10)) & ~ (e12 = op1(e10, op1(e10, e12)))) | ~ sP2), inference(usedef, [], [e23])).
fof(e23, plain, (sP2 <=> (~ (e10 = op1(op1(e10, e12), e12)) & (e12 = op1(op1(e10, e12), e10)) & ~ (e12 = op1(e10, op1(e10, e12))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f24, plain, ((~ (e10 = op1(op1(e10, e13), e13)) & (e13 = op1(op1(e10, e13), e10)) & ~ (e13 = op1(e10, op1(e10, e13)))) | ~ sP3), inference(usedef, [], [e24])).
fof(e24, plain, (sP3 <=> (~ (e10 = op1(op1(e10, e13), e13)) & (e13 = op1(op1(e10, e13), e10)) & ~ (e13 = op1(e10, op1(e10, e13))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f25, plain, ((~ (e11 = op1(op1(e11, e10), e10)) & (e10 = op1(op1(e11, e10), e11)) & ~ (e10 = op1(e11, op1(e11, e10)))) | ~ sP4), inference(usedef, [], [e25])).
fof(e25, plain, (sP4 <=> (~ (e11 = op1(op1(e11, e10), e10)) & (e10 = op1(op1(e11, e10), e11)) & ~ (e10 = op1(e11, op1(e11, e10))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f26, plain, ((~ (e11 = op1(op1(e11, e11), e11)) & (e11 = op1(op1(e11, e11), e11)) & ~ (e11 = op1(e11, op1(e11, e11)))) | ~ sP5), inference(usedef, [], [e26])).
fof(e26, plain, (sP5 <=> (~ (e11 = op1(op1(e11, e11), e11)) & (e11 = op1(op1(e11, e11), e11)) & ~ (e11 = op1(e11, op1(e11, e11))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f27, plain, ((~ (e11 = op1(op1(e11, e12), e12)) & (e12 = op1(op1(e11, e12), e11)) & ~ (e12 = op1(e11, op1(e11, e12)))) | ~ sP6), inference(usedef, [], [e27])).
fof(e27, plain, (sP6 <=> (~ (e11 = op1(op1(e11, e12), e12)) & (e12 = op1(op1(e11, e12), e11)) & ~ (e12 = op1(e11, op1(e11, e12))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f28, plain, ((~ (e11 = op1(op1(e11, e13), e13)) & (e13 = op1(op1(e11, e13), e11)) & ~ (e13 = op1(e11, op1(e11, e13)))) | ~ sP7), inference(usedef, [], [e28])).
fof(e28, plain, (sP7 <=> (~ (e11 = op1(op1(e11, e13), e13)) & (e13 = op1(op1(e11, e13), e11)) & ~ (e13 = op1(e11, op1(e11, e13))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f29, plain, ((~ (e12 = op1(op1(e12, e10), e10)) & (e10 = op1(op1(e12, e10), e12)) & ~ (e10 = op1(e12, op1(e12, e10)))) | ~ sP8), inference(usedef, [], [e29])).
fof(e29, plain, (sP8 <=> (~ (e12 = op1(op1(e12, e10), e10)) & (e10 = op1(op1(e12, e10), e12)) & ~ (e10 = op1(e12, op1(e12, e10))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f30, plain, ((~ (e12 = op1(op1(e12, e11), e11)) & (e11 = op1(op1(e12, e11), e12)) & ~ (e11 = op1(e12, op1(e12, e11)))) | ~ sP9), inference(usedef, [], [e30])).
fof(e30, plain, (sP9 <=> (~ (e12 = op1(op1(e12, e11), e11)) & (e11 = op1(op1(e12, e11), e12)) & ~ (e11 = op1(e12, op1(e12, e11))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f31, plain, ((~ (e12 = op1(op1(e12, e12), e12)) & (e12 = op1(op1(e12, e12), e12)) & ~ (e12 = op1(e12, op1(e12, e12)))) | ~ sP10), inference(usedef, [], [e31])).
fof(e31, plain, (sP10 <=> (~ (e12 = op1(op1(e12, e12), e12)) & (e12 = op1(op1(e12, e12), e12)) & ~ (e12 = op1(e12, op1(e12, e12))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f32, plain, ((~ (e12 = op1(op1(e12, e13), e13)) & (e13 = op1(op1(e12, e13), e12)) & ~ (e13 = op1(e12, op1(e12, e13)))) | ~ sP11), inference(usedef, [], [e32])).
fof(e32, plain, (sP11 <=> (~ (e12 = op1(op1(e12, e13), e13)) & (e13 = op1(op1(e12, e13), e12)) & ~ (e13 = op1(e12, op1(e12, e13))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f33, plain, ((~ (e13 = op1(op1(e13, e10), e10)) & (e10 = op1(op1(e13, e10), e13)) & ~ (e10 = op1(e13, op1(e13, e10)))) | ~ sP12), inference(usedef, [], [e33])).
fof(e33, plain, (sP12 <=> (~ (e13 = op1(op1(e13, e10), e10)) & (e10 = op1(op1(e13, e10), e13)) & ~ (e10 = op1(e13, op1(e13, e10))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f34, plain, ((~ (e13 = op1(op1(e13, e11), e11)) & (e11 = op1(op1(e13, e11), e13)) & ~ (e11 = op1(e13, op1(e13, e11)))) | ~ sP13), inference(usedef, [], [e34])).
fof(e34, plain, (sP13 <=> (~ (e13 = op1(op1(e13, e11), e11)) & (e11 = op1(op1(e13, e11), e13)) & ~ (e11 = op1(e13, op1(e13, e11))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f35, plain, ((~ (e13 = op1(op1(e13, e12), e12)) & (e12 = op1(op1(e13, e12), e13)) & ~ (e12 = op1(e13, op1(e13, e12)))) | ~ sP14), inference(usedef, [], [e35])).
fof(e35, plain, (sP14 <=> (~ (e13 = op1(op1(e13, e12), e12)) & (e12 = op1(op1(e13, e12), e13)) & ~ (e12 = op1(e13, op1(e13, e12))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f36, plain, ((~ (e13 = op1(e13, op1(e10, e10))) & ~ (e12 = op1(e12, op1(e10, e10))) & ~ (e11 = op1(e11, op1(e10, e10))) & ~ (e10 = op1(e10, op1(e10, e10))) & (op1(e10, e10) = op1(e10, e10))) | ~ sP15), inference(usedef, [], [e36])).
fof(e36, plain, (sP15 <=> (~ (e13 = op1(e13, op1(e10, e10))) & ~ (e12 = op1(e12, op1(e10, e10))) & ~ (e11 = op1(e11, op1(e10, e10))) & ~ (e10 = op1(e10, op1(e10, e10))) & (op1(e10, e10) = op1(e10, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f37, plain, ((~ (e13 = op1(e13, op1(e10, e11))) & ~ (e12 = op1(e12, op1(e10, e11))) & ~ (e11 = op1(e11, op1(e10, e11))) & ~ (e10 = op1(e10, op1(e10, e11))) & (op1(e10, e11) = op1(e11, e10))) | ~ sP16), inference(usedef, [], [e37])).
fof(e37, plain, (sP16 <=> (~ (e13 = op1(e13, op1(e10, e11))) & ~ (e12 = op1(e12, op1(e10, e11))) & ~ (e11 = op1(e11, op1(e10, e11))) & ~ (e10 = op1(e10, op1(e10, e11))) & (op1(e10, e11) = op1(e11, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f38, plain, ((~ (e13 = op1(e13, op1(e10, e12))) & ~ (e12 = op1(e12, op1(e10, e12))) & ~ (e11 = op1(e11, op1(e10, e12))) & ~ (e10 = op1(e10, op1(e10, e12))) & (op1(e10, e12) = op1(e12, e10))) | ~ sP17), inference(usedef, [], [e38])).
fof(e38, plain, (sP17 <=> (~ (e13 = op1(e13, op1(e10, e12))) & ~ (e12 = op1(e12, op1(e10, e12))) & ~ (e11 = op1(e11, op1(e10, e12))) & ~ (e10 = op1(e10, op1(e10, e12))) & (op1(e10, e12) = op1(e12, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f39, plain, ((~ (e13 = op1(e13, op1(e10, e13))) & ~ (e12 = op1(e12, op1(e10, e13))) & ~ (e11 = op1(e11, op1(e10, e13))) & ~ (e10 = op1(e10, op1(e10, e13))) & (op1(e10, e13) = op1(e13, e10))) | ~ sP18), inference(usedef, [], [e39])).
fof(e39, plain, (sP18 <=> (~ (e13 = op1(e13, op1(e10, e13))) & ~ (e12 = op1(e12, op1(e10, e13))) & ~ (e11 = op1(e11, op1(e10, e13))) & ~ (e10 = op1(e10, op1(e10, e13))) & (op1(e10, e13) = op1(e13, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f40, plain, ((~ (e13 = op1(e13, op1(e11, e10))) & ~ (e12 = op1(e12, op1(e11, e10))) & ~ (e11 = op1(e11, op1(e11, e10))) & ~ (e10 = op1(e10, op1(e11, e10))) & (op1(e10, e11) = op1(e11, e10))) | ~ sP19), inference(usedef, [], [e40])).
fof(e40, plain, (sP19 <=> (~ (e13 = op1(e13, op1(e11, e10))) & ~ (e12 = op1(e12, op1(e11, e10))) & ~ (e11 = op1(e11, op1(e11, e10))) & ~ (e10 = op1(e10, op1(e11, e10))) & (op1(e10, e11) = op1(e11, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP19])])).
fof(f41, plain, ((~ (e13 = op1(e13, op1(e11, e11))) & ~ (e12 = op1(e12, op1(e11, e11))) & ~ (e11 = op1(e11, op1(e11, e11))) & ~ (e10 = op1(e10, op1(e11, e11))) & (op1(e11, e11) = op1(e11, e11))) | ~ sP20), inference(usedef, [], [e41])).
fof(e41, plain, (sP20 <=> (~ (e13 = op1(e13, op1(e11, e11))) & ~ (e12 = op1(e12, op1(e11, e11))) & ~ (e11 = op1(e11, op1(e11, e11))) & ~ (e10 = op1(e10, op1(e11, e11))) & (op1(e11, e11) = op1(e11, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP20])])).
fof(f42, plain, ((~ (e13 = op1(e13, op1(e11, e12))) & ~ (e12 = op1(e12, op1(e11, e12))) & ~ (e11 = op1(e11, op1(e11, e12))) & ~ (e10 = op1(e10, op1(e11, e12))) & (op1(e11, e12) = op1(e12, e11))) | ~ sP21), inference(usedef, [], [e42])).
fof(e42, plain, (sP21 <=> (~ (e13 = op1(e13, op1(e11, e12))) & ~ (e12 = op1(e12, op1(e11, e12))) & ~ (e11 = op1(e11, op1(e11, e12))) & ~ (e10 = op1(e10, op1(e11, e12))) & (op1(e11, e12) = op1(e12, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP21])])).
fof(f43, plain, ((~ (e13 = op1(e13, op1(e11, e13))) & ~ (e12 = op1(e12, op1(e11, e13))) & ~ (e11 = op1(e11, op1(e11, e13))) & ~ (e10 = op1(e10, op1(e11, e13))) & (op1(e11, e13) = op1(e13, e11))) | ~ sP22), inference(usedef, [], [e43])).
fof(e43, plain, (sP22 <=> (~ (e13 = op1(e13, op1(e11, e13))) & ~ (e12 = op1(e12, op1(e11, e13))) & ~ (e11 = op1(e11, op1(e11, e13))) & ~ (e10 = op1(e10, op1(e11, e13))) & (op1(e11, e13) = op1(e13, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP22])])).
fof(f44, plain, ((~ (e13 = op1(e13, op1(e12, e10))) & ~ (e12 = op1(e12, op1(e12, e10))) & ~ (e11 = op1(e11, op1(e12, e10))) & ~ (e10 = op1(e10, op1(e12, e10))) & (op1(e10, e12) = op1(e12, e10))) | ~ sP23), inference(usedef, [], [e44])).
fof(e44, plain, (sP23 <=> (~ (e13 = op1(e13, op1(e12, e10))) & ~ (e12 = op1(e12, op1(e12, e10))) & ~ (e11 = op1(e11, op1(e12, e10))) & ~ (e10 = op1(e10, op1(e12, e10))) & (op1(e10, e12) = op1(e12, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP23])])).
fof(f45, plain, ((~ (e13 = op1(e13, op1(e12, e11))) & ~ (e12 = op1(e12, op1(e12, e11))) & ~ (e11 = op1(e11, op1(e12, e11))) & ~ (e10 = op1(e10, op1(e12, e11))) & (op1(e11, e12) = op1(e12, e11))) | ~ sP24), inference(usedef, [], [e45])).
fof(e45, plain, (sP24 <=> (~ (e13 = op1(e13, op1(e12, e11))) & ~ (e12 = op1(e12, op1(e12, e11))) & ~ (e11 = op1(e11, op1(e12, e11))) & ~ (e10 = op1(e10, op1(e12, e11))) & (op1(e11, e12) = op1(e12, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP24])])).
fof(f46, plain, ((~ (e13 = op1(e13, op1(e12, e12))) & ~ (e12 = op1(e12, op1(e12, e12))) & ~ (e11 = op1(e11, op1(e12, e12))) & ~ (e10 = op1(e10, op1(e12, e12))) & (op1(e12, e12) = op1(e12, e12))) | ~ sP25), inference(usedef, [], [e46])).
fof(e46, plain, (sP25 <=> (~ (e13 = op1(e13, op1(e12, e12))) & ~ (e12 = op1(e12, op1(e12, e12))) & ~ (e11 = op1(e11, op1(e12, e12))) & ~ (e10 = op1(e10, op1(e12, e12))) & (op1(e12, e12) = op1(e12, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP25])])).
fof(f47, plain, ((~ (e13 = op1(e13, op1(e12, e13))) & ~ (e12 = op1(e12, op1(e12, e13))) & ~ (e11 = op1(e11, op1(e12, e13))) & ~ (e10 = op1(e10, op1(e12, e13))) & (op1(e12, e13) = op1(e13, e12))) | ~ sP26), inference(usedef, [], [e47])).
fof(e47, plain, (sP26 <=> (~ (e13 = op1(e13, op1(e12, e13))) & ~ (e12 = op1(e12, op1(e12, e13))) & ~ (e11 = op1(e11, op1(e12, e13))) & ~ (e10 = op1(e10, op1(e12, e13))) & (op1(e12, e13) = op1(e13, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP26])])).
fof(f48, plain, ((~ (e13 = op1(e13, op1(e13, e10))) & ~ (e12 = op1(e12, op1(e13, e10))) & ~ (e11 = op1(e11, op1(e13, e10))) & ~ (e10 = op1(e10, op1(e13, e10))) & (op1(e10, e13) = op1(e13, e10))) | ~ sP27), inference(usedef, [], [e48])).
fof(e48, plain, (sP27 <=> (~ (e13 = op1(e13, op1(e13, e10))) & ~ (e12 = op1(e12, op1(e13, e10))) & ~ (e11 = op1(e11, op1(e13, e10))) & ~ (e10 = op1(e10, op1(e13, e10))) & (op1(e10, e13) = op1(e13, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP27])])).
fof(f49, plain, ((~ (e13 = op1(e13, op1(e13, e11))) & ~ (e12 = op1(e12, op1(e13, e11))) & ~ (e11 = op1(e11, op1(e13, e11))) & ~ (e10 = op1(e10, op1(e13, e11))) & (op1(e11, e13) = op1(e13, e11))) | ~ sP28), inference(usedef, [], [e49])).
fof(e49, plain, (sP28 <=> (~ (e13 = op1(e13, op1(e13, e11))) & ~ (e12 = op1(e12, op1(e13, e11))) & ~ (e11 = op1(e11, op1(e13, e11))) & ~ (e10 = op1(e10, op1(e13, e11))) & (op1(e11, e13) = op1(e13, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP28])])).
fof(f50, plain, ((~ (e13 = op1(e13, op1(e13, e12))) & ~ (e12 = op1(e12, op1(e13, e12))) & ~ (e11 = op1(e11, op1(e13, e12))) & ~ (e10 = op1(e10, op1(e13, e12))) & (op1(e12, e13) = op1(e13, e12))) | ~ sP29), inference(usedef, [], [e50])).
fof(e50, plain, (sP29 <=> (~ (e13 = op1(e13, op1(e13, e12))) & ~ (e12 = op1(e12, op1(e13, e12))) & ~ (e11 = op1(e11, op1(e13, e12))) & ~ (e10 = op1(e10, op1(e13, e12))) & (op1(e12, e13) = op1(e13, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP29])])).
fof(f10, plain, (((~ (e13 = op1(e13, op1(e13, e13))) & ~ (e12 = op1(e12, op1(e13, e13))) & ~ (e11 = op1(e11, op1(e13, e13))) & ~ (e10 = op1(e10, op1(e13, e13))) & (op1(e13, e13) = op1(e13, e13))) | (~ (e13 = op1(e13, op1(e13, e12))) & ~ (e12 = op1(e12, op1(e13, e12))) & ~ (e11 = op1(e11, op1(e13, e12))) & ~ (e10 = op1(e10, op1(e13, e12))) & (op1(e12, e13) = op1(e13, e12))) | (~ (e13 = op1(e13, op1(e13, e11))) & ~ (e12 = op1(e12, op1(e13, e11))) & ~ (e11 = op1(e11, op1(e13, e11))) & ~ (e10 = op1(e10, op1(e13, e11))) & (op1(e11, e13) = op1(e13, e11))) | (~ (e13 = op1(e13, op1(e13, e10))) & ~ (e12 = op1(e12, op1(e13, e10))) & ~ (e11 = op1(e11, op1(e13, e10))) & ~ (e10 = op1(e10, op1(e13, e10))) & (op1(e10, e13) = op1(e13, e10))) | (~ (e13 = op1(e13, op1(e12, e13))) & ~ (e12 = op1(e12, op1(e12, e13))) & ~ (e11 = op1(e11, op1(e12, e13))) & ~ (e10 = op1(e10, op1(e12, e13))) & (op1(e12, e13) = op1(e13, e12))) | (~ (e13 = op1(e13, op1(e12, e12))) & ~ (e12 = op1(e12, op1(e12, e12))) & ~ (e11 = op1(e11, op1(e12, e12))) & ~ (e10 = op1(e10, op1(e12, e12))) & (op1(e12, e12) = op1(e12, e12))) | (~ (e13 = op1(e13, op1(e12, e11))) & ~ (e12 = op1(e12, op1(e12, e11))) & ~ (e11 = op1(e11, op1(e12, e11))) & ~ (e10 = op1(e10, op1(e12, e11))) & (op1(e11, e12) = op1(e12, e11))) | (~ (e13 = op1(e13, op1(e12, e10))) & ~ (e12 = op1(e12, op1(e12, e10))) & ~ (e11 = op1(e11, op1(e12, e10))) & ~ (e10 = op1(e10, op1(e12, e10))) & (op1(e10, e12) = op1(e12, e10))) | (~ (e13 = op1(e13, op1(e11, e13))) & ~ (e12 = op1(e12, op1(e11, e13))) & ~ (e11 = op1(e11, op1(e11, e13))) & ~ (e10 = op1(e10, op1(e11, e13))) & (op1(e11, e13) = op1(e13, e11))) | (~ (e13 = op1(e13, op1(e11, e12))) & ~ (e12 = op1(e12, op1(e11, e12))) & ~ (e11 = op1(e11, op1(e11, e12))) & ~ (e10 = op1(e10, op1(e11, e12))) & (op1(e11, e12) = op1(e12, e11))) | (~ (e13 = op1(e13, op1(e11, e11))) & ~ (e12 = op1(e12, op1(e11, e11))) & ~ (e11 = op1(e11, op1(e11, e11))) & ~ (e10 = op1(e10, op1(e11, e11))) & (op1(e11, e11) = op1(e11, e11))) | (~ (e13 = op1(e13, op1(e11, e10))) & ~ (e12 = op1(e12, op1(e11, e10))) & ~ (e11 = op1(e11, op1(e11, e10))) & ~ (e10 = op1(e10, op1(e11, e10))) & (op1(e10, e11) = op1(e11, e10))) | (~ (e13 = op1(e13, op1(e10, e13))) & ~ (e12 = op1(e12, op1(e10, e13))) & ~ (e11 = op1(e11, op1(e10, e13))) & ~ (e10 = op1(e10, op1(e10, e13))) & (op1(e10, e13) = op1(e13, e10))) | (~ (e13 = op1(e13, op1(e10, e12))) & ~ (e12 = op1(e12, op1(e10, e12))) & ~ (e11 = op1(e11, op1(e10, e12))) & ~ (e10 = op1(e10, op1(e10, e12))) & (op1(e10, e12) = op1(e12, e10))) | (~ (e13 = op1(e13, op1(e10, e11))) & ~ (e12 = op1(e12, op1(e10, e11))) & ~ (e11 = op1(e11, op1(e10, e11))) & ~ (e10 = op1(e10, op1(e10, e11))) & (op1(e10, e11) = op1(e11, e10))) | (~ (e13 = op1(e13, op1(e10, e10))) & ~ (e12 = op1(e12, op1(e10, e10))) & ~ (e11 = op1(e11, op1(e10, e10))) & ~ (e10 = op1(e10, op1(e10, e10))) & (op1(e10, e10) = op1(e10, e10)))) & ((~ (e13 = op1(op1(e13, e13), e13)) & (e13 = op1(op1(e13, e13), e13)) & ~ (e13 = op1(e13, op1(e13, e13)))) | (~ (e13 = op1(op1(e13, e12), e12)) & (e12 = op1(op1(e13, e12), e13)) & ~ (e12 = op1(e13, op1(e13, e12)))) | (~ (e13 = op1(op1(e13, e11), e11)) & (e11 = op1(op1(e13, e11), e13)) & ~ (e11 = op1(e13, op1(e13, e11)))) | (~ (e13 = op1(op1(e13, e10), e10)) & (e10 = op1(op1(e13, e10), e13)) & ~ (e10 = op1(e13, op1(e13, e10)))) | (~ (e12 = op1(op1(e12, e13), e13)) & (e13 = op1(op1(e12, e13), e12)) & ~ (e13 = op1(e12, op1(e12, e13)))) | (~ (e12 = op1(op1(e12, e12), e12)) & (e12 = op1(op1(e12, e12), e12)) & ~ (e12 = op1(e12, op1(e12, e12)))) | (~ (e12 = op1(op1(e12, e11), e11)) & (e11 = op1(op1(e12, e11), e12)) & ~ (e11 = op1(e12, op1(e12, e11)))) | (~ (e12 = op1(op1(e12, e10), e10)) & (e10 = op1(op1(e12, e10), e12)) & ~ (e10 = op1(e12, op1(e12, e10)))) | (~ (e11 = op1(op1(e11, e13), e13)) & (e13 = op1(op1(e11, e13), e11)) & ~ (e13 = op1(e11, op1(e11, e13)))) | (~ (e11 = op1(op1(e11, e12), e12)) & (e12 = op1(op1(e11, e12), e11)) & ~ (e12 = op1(e11, op1(e11, e12)))) | (~ (e11 = op1(op1(e11, e11), e11)) & (e11 = op1(op1(e11, e11), e11)) & ~ (e11 = op1(e11, op1(e11, e11)))) | (~ (e11 = op1(op1(e11, e10), e10)) & (e10 = op1(op1(e11, e10), e11)) & ~ (e10 = op1(e11, op1(e11, e10)))) | (~ (e10 = op1(op1(e10, e13), e13)) & (e13 = op1(op1(e10, e13), e10)) & ~ (e13 = op1(e10, op1(e10, e13)))) | (~ (e10 = op1(op1(e10, e12), e12)) & (e12 = op1(op1(e10, e12), e10)) & ~ (e12 = op1(e10, op1(e10, e12)))) | (~ (e10 = op1(op1(e10, e11), e11)) & (e11 = op1(op1(e10, e11), e10)) & ~ (e11 = op1(e10, op1(e10, e11)))) | (~ (e10 = op1(op1(e10, e10), e10)) & (e10 = op1(op1(e10, e10), e10)) & ~ (e10 = op1(e10, op1(e10, e10)))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG104+1.p', ax10)).
fof(f2020, plain, (spl72_262 | spl72_258 | spl72_254 | spl72_250 | spl72_246 | spl72_244 | spl72_240 | spl72_236 | spl72_232 | spl72_228 | spl72_226 | spl72_222 | spl72_218 | spl72_214 | spl72_210 | spl72_268), inference(avatar_split_clause, [], [f509, f2016, f1735, f1754, f1773, f1792, f1811, f1822, f1841, f1860, f1879, f1898, f1909, f1928, f1947, f1966, f1985])).
fof(f2016, plain, (spl72_268 <=> (e13 = op1(op1(e13, e13), e13))), introduced(avatar_definition, [new_symbols(naming, [spl72_268])])).
fof(f509, plain, ((e13 = op1(op1(e13, e13), e13)) | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f51])).
fof(f2019, plain, (spl72_262 | spl72_258 | spl72_254 | spl72_250 | spl72_246 | spl72_244 | spl72_240 | spl72_236 | spl72_232 | spl72_228 | spl72_226 | spl72_222 | spl72_218 | spl72_214 | spl72_210 | ~ spl72_268), inference(avatar_split_clause, [], [f510, f2016, f1735, f1754, f1773, f1792, f1811, f1822, f1841, f1860, f1879, f1898, f1909, f1928, f1947, f1966, f1985])).
fof(f510, plain, (~ (e13 = op1(op1(e13, e13), e13)) | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f51])).
fof(f2014, plain, (spl72_205 | spl72_200 | spl72_195 | spl72_190 | spl72_184 | spl72_179 | spl72_174 | spl72_169 | spl72_163 | spl72_157 | spl72_152 | spl72_147 | spl72_141 | spl72_135 | spl72_129 | ~ spl72_267), inference(avatar_split_clause, [], [f512, f2011, f1339, f1368, f1397, f1426, f1451, f1475, f1504, f1533, f1558, f1583, f1607, f1636, f1661, f1686, f1711])).
fof(f1711, plain, (spl72_205 <=> sP15), introduced(avatar_definition, [new_symbols(naming, [spl72_205])])).
fof(f1686, plain, (spl72_200 <=> sP16), introduced(avatar_definition, [new_symbols(naming, [spl72_200])])).
fof(f1661, plain, (spl72_195 <=> sP17), introduced(avatar_definition, [new_symbols(naming, [spl72_195])])).
fof(f1636, plain, (spl72_190 <=> sP18), introduced(avatar_definition, [new_symbols(naming, [spl72_190])])).
fof(f1607, plain, (spl72_184 <=> sP19), introduced(avatar_definition, [new_symbols(naming, [spl72_184])])).
fof(f1583, plain, (spl72_179 <=> sP20), introduced(avatar_definition, [new_symbols(naming, [spl72_179])])).
fof(f1558, plain, (spl72_174 <=> sP21), introduced(avatar_definition, [new_symbols(naming, [spl72_174])])).
fof(f1533, plain, (spl72_169 <=> sP22), introduced(avatar_definition, [new_symbols(naming, [spl72_169])])).
fof(f1504, plain, (spl72_163 <=> sP23), introduced(avatar_definition, [new_symbols(naming, [spl72_163])])).
fof(f1475, plain, (spl72_157 <=> sP24), introduced(avatar_definition, [new_symbols(naming, [spl72_157])])).
fof(f1451, plain, (spl72_152 <=> sP25), introduced(avatar_definition, [new_symbols(naming, [spl72_152])])).
fof(f1426, plain, (spl72_147 <=> sP26), introduced(avatar_definition, [new_symbols(naming, [spl72_147])])).
fof(f1397, plain, (spl72_141 <=> sP27), introduced(avatar_definition, [new_symbols(naming, [spl72_141])])).
fof(f1368, plain, (spl72_135 <=> sP28), introduced(avatar_definition, [new_symbols(naming, [spl72_135])])).
fof(f1339, plain, (spl72_129 <=> sP29), introduced(avatar_definition, [new_symbols(naming, [spl72_129])])).
fof(f512, plain, (~ (e10 = op1(e10, op1(e13, e13))) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15), inference(cnf_transformation, [], [f51])).
fof(f2004, plain, (spl72_205 | spl72_200 | spl72_195 | spl72_190 | spl72_184 | spl72_179 | spl72_174 | spl72_169 | spl72_163 | spl72_157 | spl72_152 | spl72_147 | spl72_141 | spl72_135 | spl72_129 | ~ spl72_265), inference(avatar_split_clause, [], [f514, f2001, f1339, f1368, f1397, f1426, f1451, f1475, f1504, f1533, f1558, f1583, f1607, f1636, f1661, f1686, f1711])).
fof(f514, plain, (~ (e12 = op1(e12, op1(e13, e13))) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15), inference(cnf_transformation, [], [f51])).
fof(f1993, plain, (~ spl72_262 | spl72_263), inference(avatar_split_clause, [], [f506, f1989, f1985])).
fof(f1989, plain, (spl72_263 <=> (e10 = op1(op1(e10, e10), e10))), introduced(avatar_definition, [new_symbols(naming, [spl72_263])])).
fof(f506, plain, ((e10 = op1(op1(e10, e10), e10)) | ~ sP0), inference(cnf_transformation, [], [f125])).
fof(f125, plain, ((~ (e10 = op1(op1(e10, e10), e10)) & (e10 = op1(op1(e10, e10), e10)) & ~ (e10 = op1(e10, op1(e10, e10)))) | ~ sP0), inference(nnf_transformation, [], [f21])).
fof(f1992, plain, (~ spl72_262 | ~ spl72_263), inference(avatar_split_clause, [], [f507, f1989, f1985])).
fof(f507, plain, (~ (e10 = op1(op1(e10, e10), e10)) | ~ sP0), inference(cnf_transformation, [], [f125])).
fof(f1978, plain, (~ spl72_258 | spl72_260), inference(avatar_split_clause, [], [f503, f1975, f1966])).
fof(f503, plain, ((e11 = op1(op1(e10, e11), e10)) | ~ sP1), inference(cnf_transformation, [], [f124])).
fof(f124, plain, ((~ (e10 = op1(op1(e10, e11), e11)) & (e11 = op1(op1(e10, e11), e10)) & ~ (e11 = op1(e10, op1(e10, e11)))) | ~ sP1), inference(nnf_transformation, [], [f22])).
fof(f1973, plain, (~ spl72_258 | ~ spl72_259), inference(avatar_split_clause, [], [f504, f1970, f1966])).
fof(f504, plain, (~ (e10 = op1(op1(e10, e11), e11)) | ~ sP1), inference(cnf_transformation, [], [f124])).
fof(f1959, plain, (~ spl72_254 | spl72_256), inference(avatar_split_clause, [], [f500, f1956, f1947])).
fof(f500, plain, ((e12 = op1(op1(e10, e12), e10)) | ~ sP2), inference(cnf_transformation, [], [f123])).
fof(f123, plain, ((~ (e10 = op1(op1(e10, e12), e12)) & (e12 = op1(op1(e10, e12), e10)) & ~ (e12 = op1(e10, op1(e10, e12)))) | ~ sP2), inference(nnf_transformation, [], [f23])).
fof(f1935, plain, (~ spl72_250 | ~ spl72_251), inference(avatar_split_clause, [], [f498, f1932, f1928])).
fof(f498, plain, (~ (e10 = op1(op1(e10, e13), e13)) | ~ sP3), inference(cnf_transformation, [], [f122])).
fof(f122, plain, ((~ (e10 = op1(op1(e10, e13), e13)) & (e13 = op1(op1(e10, e13), e10)) & ~ (e13 = op1(e10, op1(e10, e13)))) | ~ sP3), inference(nnf_transformation, [], [f24])).
fof(f1926, plain, (~ spl72_246 | ~ spl72_249), inference(avatar_split_clause, [], [f493, f1923, f1909])).
fof(f493, plain, (~ (e10 = op1(e11, op1(e11, e10))) | ~ sP4), inference(cnf_transformation, [], [f121])).
fof(f121, plain, ((~ (e11 = op1(op1(e11, e10), e10)) & (e10 = op1(op1(e11, e10), e11)) & ~ (e10 = op1(e11, op1(e11, e10)))) | ~ sP4), inference(nnf_transformation, [], [f25])).
fof(f1921, plain, (~ spl72_246 | spl72_248), inference(avatar_split_clause, [], [f494, f1918, f1909])).
fof(f494, plain, ((e10 = op1(op1(e11, e10), e11)) | ~ sP4), inference(cnf_transformation, [], [f121])).
fof(f1916, plain, (~ spl72_246 | ~ spl72_247), inference(avatar_split_clause, [], [f495, f1913, f1909])).
fof(f495, plain, (~ (e11 = op1(op1(e11, e10), e10)) | ~ sP4), inference(cnf_transformation, [], [f121])).
fof(f1907, plain, (~ spl72_244 | ~ spl72_182), inference(avatar_split_clause, [], [f490, f1597, f1898])).
fof(f490, plain, (~ (e11 = op1(e11, op1(e11, e11))) | ~ sP5), inference(cnf_transformation, [], [f120])).
fof(f120, plain, ((~ (e11 = op1(op1(e11, e11), e11)) & (e11 = op1(op1(e11, e11), e11)) & ~ (e11 = op1(e11, op1(e11, e11)))) | ~ sP5), inference(nnf_transformation, [], [f26])).
fof(f1906, plain, (~ spl72_244 | spl72_245), inference(avatar_split_clause, [], [f491, f1902, f1898])).
fof(f1902, plain, (spl72_245 <=> (e11 = op1(op1(e11, e11), e11))), introduced(avatar_definition, [new_symbols(naming, [spl72_245])])).
fof(f491, plain, ((e11 = op1(op1(e11, e11), e11)) | ~ sP5), inference(cnf_transformation, [], [f120])).
fof(f1905, plain, (~ spl72_244 | ~ spl72_245), inference(avatar_split_clause, [], [f492, f1902, f1898])).
fof(f492, plain, (~ (e11 = op1(op1(e11, e11), e11)) | ~ sP5), inference(cnf_transformation, [], [f120])).
fof(f1896, plain, (~ spl72_240 | ~ spl72_243), inference(avatar_split_clause, [], [f487, f1893, f1879])).
fof(f487, plain, (~ (e12 = op1(e11, op1(e11, e12))) | ~ sP6), inference(cnf_transformation, [], [f119])).
fof(f119, plain, ((~ (e11 = op1(op1(e11, e12), e12)) & (e12 = op1(op1(e11, e12), e11)) & ~ (e12 = op1(e11, op1(e11, e12)))) | ~ sP6), inference(nnf_transformation, [], [f27])).
fof(f1891, plain, (~ spl72_240 | spl72_242), inference(avatar_split_clause, [], [f488, f1888, f1879])).
fof(f488, plain, ((e12 = op1(op1(e11, e12), e11)) | ~ sP6), inference(cnf_transformation, [], [f119])).
fof(f1886, plain, (~ spl72_240 | ~ spl72_241), inference(avatar_split_clause, [], [f489, f1883, f1879])).
fof(f489, plain, (~ (e11 = op1(op1(e11, e12), e12)) | ~ sP6), inference(cnf_transformation, [], [f119])).
fof(f1877, plain, (~ spl72_236 | ~ spl72_239), inference(avatar_split_clause, [], [f484, f1874, f1860])).
fof(f484, plain, (~ (e13 = op1(e11, op1(e11, e13))) | ~ sP7), inference(cnf_transformation, [], [f118])).
fof(f118, plain, ((~ (e11 = op1(op1(e11, e13), e13)) & (e13 = op1(op1(e11, e13), e11)) & ~ (e13 = op1(e11, op1(e11, e13)))) | ~ sP7), inference(nnf_transformation, [], [f28])).
fof(f1872, plain, (~ spl72_236 | spl72_238), inference(avatar_split_clause, [], [f485, f1869, f1860])).
fof(f485, plain, ((e13 = op1(op1(e11, e13), e11)) | ~ sP7), inference(cnf_transformation, [], [f118])).
fof(f1867, plain, (~ spl72_236 | ~ spl72_237), inference(avatar_split_clause, [], [f486, f1864, f1860])).
fof(f486, plain, (~ (e11 = op1(op1(e11, e13), e13)) | ~ sP7), inference(cnf_transformation, [], [f118])).
fof(f1858, plain, (~ spl72_232 | ~ spl72_235), inference(avatar_split_clause, [], [f481, f1855, f1841])).
fof(f481, plain, (~ (e10 = op1(e12, op1(e12, e10))) | ~ sP8), inference(cnf_transformation, [], [f117])).
fof(f117, plain, ((~ (e12 = op1(op1(e12, e10), e10)) & (e10 = op1(op1(e12, e10), e12)) & ~ (e10 = op1(e12, op1(e12, e10)))) | ~ sP8), inference(nnf_transformation, [], [f29])).
fof(f1853, plain, (~ spl72_232 | spl72_234), inference(avatar_split_clause, [], [f482, f1850, f1841])).
fof(f482, plain, ((e10 = op1(op1(e12, e10), e12)) | ~ sP8), inference(cnf_transformation, [], [f117])).
fof(f1839, plain, (~ spl72_228 | ~ spl72_231), inference(avatar_split_clause, [], [f478, f1836, f1822])).
fof(f478, plain, (~ (e11 = op1(e12, op1(e12, e11))) | ~ sP9), inference(cnf_transformation, [], [f116])).
fof(f116, plain, ((~ (e12 = op1(op1(e12, e11), e11)) & (e11 = op1(op1(e12, e11), e12)) & ~ (e11 = op1(e12, op1(e12, e11)))) | ~ sP9), inference(nnf_transformation, [], [f30])).
fof(f1829, plain, (~ spl72_228 | ~ spl72_229), inference(avatar_split_clause, [], [f480, f1826, f1822])).
fof(f480, plain, (~ (e12 = op1(op1(e12, e11), e11)) | ~ sP9), inference(cnf_transformation, [], [f116])).
fof(f1819, plain, (~ spl72_226 | spl72_227), inference(avatar_split_clause, [], [f476, f1815, f1811])).
fof(f1815, plain, (spl72_227 <=> (e12 = op1(op1(e12, e12), e12))), introduced(avatar_definition, [new_symbols(naming, [spl72_227])])).
fof(f476, plain, ((e12 = op1(op1(e12, e12), e12)) | ~ sP10), inference(cnf_transformation, [], [f115])).
fof(f115, plain, ((~ (e12 = op1(op1(e12, e12), e12)) & (e12 = op1(op1(e12, e12), e12)) & ~ (e12 = op1(e12, op1(e12, e12)))) | ~ sP10), inference(nnf_transformation, [], [f31])).
fof(f1818, plain, (~ spl72_226 | ~ spl72_227), inference(avatar_split_clause, [], [f477, f1815, f1811])).
fof(f477, plain, (~ (e12 = op1(op1(e12, e12), e12)) | ~ sP10), inference(cnf_transformation, [], [f115])).
fof(f1809, plain, (~ spl72_222 | ~ spl72_225), inference(avatar_split_clause, [], [f472, f1806, f1792])).
fof(f472, plain, (~ (e13 = op1(e12, op1(e12, e13))) | ~ sP11), inference(cnf_transformation, [], [f114])).
fof(f114, plain, ((~ (e12 = op1(op1(e12, e13), e13)) & (e13 = op1(op1(e12, e13), e12)) & ~ (e13 = op1(e12, op1(e12, e13)))) | ~ sP11), inference(nnf_transformation, [], [f32])).
fof(f1804, plain, (~ spl72_222 | spl72_224), inference(avatar_split_clause, [], [f473, f1801, f1792])).
fof(f473, plain, ((e13 = op1(op1(e12, e13), e12)) | ~ sP11), inference(cnf_transformation, [], [f114])).
fof(f1799, plain, (~ spl72_222 | ~ spl72_223), inference(avatar_split_clause, [], [f474, f1796, f1792])).
fof(f474, plain, (~ (e12 = op1(op1(e12, e13), e13)) | ~ sP11), inference(cnf_transformation, [], [f114])).
fof(f1790, plain, (~ spl72_218 | ~ spl72_221), inference(avatar_split_clause, [], [f469, f1787, f1773])).
fof(f469, plain, (~ (e10 = op1(e13, op1(e13, e10))) | ~ sP12), inference(cnf_transformation, [], [f113])).
fof(f113, plain, ((~ (e13 = op1(op1(e13, e10), e10)) & (e10 = op1(op1(e13, e10), e13)) & ~ (e10 = op1(e13, op1(e13, e10)))) | ~ sP12), inference(nnf_transformation, [], [f33])).
fof(f1785, plain, (~ spl72_218 | spl72_220), inference(avatar_split_clause, [], [f470, f1782, f1773])).
fof(f470, plain, ((e10 = op1(op1(e13, e10), e13)) | ~ sP12), inference(cnf_transformation, [], [f113])).
fof(f1780, plain, (~ spl72_218 | ~ spl72_219), inference(avatar_split_clause, [], [f471, f1777, f1773])).
fof(f471, plain, (~ (e13 = op1(op1(e13, e10), e10)) | ~ sP12), inference(cnf_transformation, [], [f113])).
fof(f1771, plain, (~ spl72_214 | ~ spl72_217), inference(avatar_split_clause, [], [f466, f1768, f1754])).
fof(f466, plain, (~ (e11 = op1(e13, op1(e13, e11))) | ~ sP13), inference(cnf_transformation, [], [f112])).
fof(f112, plain, ((~ (e13 = op1(op1(e13, e11), e11)) & (e11 = op1(op1(e13, e11), e13)) & ~ (e11 = op1(e13, op1(e13, e11)))) | ~ sP13), inference(nnf_transformation, [], [f34])).
fof(f1766, plain, (~ spl72_214 | spl72_216), inference(avatar_split_clause, [], [f467, f1763, f1754])).
fof(f467, plain, ((e11 = op1(op1(e13, e11), e13)) | ~ sP13), inference(cnf_transformation, [], [f112])).
fof(f1752, plain, (~ spl72_210 | ~ spl72_213), inference(avatar_split_clause, [], [f463, f1749, f1735])).
fof(f463, plain, (~ (e12 = op1(e13, op1(e13, e12))) | ~ sP14), inference(cnf_transformation, [], [f111])).
fof(f111, plain, ((~ (e13 = op1(op1(e13, e12), e12)) & (e12 = op1(op1(e13, e12), e13)) & ~ (e12 = op1(e13, op1(e13, e12)))) | ~ sP14), inference(nnf_transformation, [], [f35])).
fof(f1747, plain, (~ spl72_210 | spl72_212), inference(avatar_split_clause, [], [f464, f1744, f1735])).
fof(f464, plain, ((e12 = op1(op1(e13, e12), e13)) | ~ sP14), inference(cnf_transformation, [], [f111])).
fof(f1742, plain, (~ spl72_210 | ~ spl72_211), inference(avatar_split_clause, [], [f465, f1739, f1735])).
fof(f465, plain, (~ (e13 = op1(op1(e13, e12), e12)) | ~ sP14), inference(cnf_transformation, [], [f111])).
fof(f1723, plain, (~ spl72_205 | ~ spl72_207), inference(avatar_split_clause, [], [f461, f1720, f1711])).
fof(f461, plain, (~ (e12 = op1(e12, op1(e10, e10))) | ~ sP15), inference(cnf_transformation, [], [f110])).
fof(f110, plain, ((~ (e13 = op1(e13, op1(e10, e10))) & ~ (e12 = op1(e12, op1(e10, e10))) & ~ (e11 = op1(e11, op1(e10, e10))) & ~ (e10 = op1(e10, op1(e10, e10))) & (op1(e10, e10) = op1(e10, e10))) | ~ sP15), inference(nnf_transformation, [], [f36])).
fof(f1708, plain, (~ spl72_200 | ~ spl72_204), inference(avatar_split_clause, [], [f454, f1705, f1686])).
fof(f454, plain, (~ (e10 = op1(e10, op1(e10, e11))) | ~ sP16), inference(cnf_transformation, [], [f109])).
fof(f109, plain, ((~ (e13 = op1(e13, op1(e10, e11))) & ~ (e12 = op1(e12, op1(e10, e11))) & ~ (e11 = op1(e11, op1(e10, e11))) & ~ (e10 = op1(e10, op1(e10, e11))) & (op1(e10, e11) = op1(e11, e10))) | ~ sP16), inference(nnf_transformation, [], [f37])).
fof(f1684, plain, (~ spl72_195 | spl72_168), inference(avatar_split_clause, [], [f448, f1528, f1661])).
fof(f448, plain, ((op1(e10, e12) = op1(e12, e10)) | ~ sP17), inference(cnf_transformation, [], [f108])).
fof(f108, plain, ((~ (e13 = op1(e13, op1(e10, e12))) & ~ (e12 = op1(e12, op1(e10, e12))) & ~ (e11 = op1(e11, op1(e10, e12))) & ~ (e10 = op1(e10, op1(e10, e12))) & (op1(e10, e12) = op1(e12, e10))) | ~ sP17), inference(nnf_transformation, [], [f38])).
fof(f1678, plain, (~ spl72_195 | ~ spl72_198), inference(avatar_split_clause, [], [f450, f1675, f1661])).
fof(f450, plain, (~ (e11 = op1(e11, op1(e10, e12))) | ~ sP17), inference(cnf_transformation, [], [f108])).
fof(f1659, plain, (~ spl72_190 | spl72_146), inference(avatar_split_clause, [], [f443, f1421, f1636])).
fof(f443, plain, ((op1(e10, e13) = op1(e13, e10)) | ~ sP18), inference(cnf_transformation, [], [f107])).
fof(f107, plain, ((~ (e13 = op1(e13, op1(e10, e13))) & ~ (e12 = op1(e12, op1(e10, e13))) & ~ (e11 = op1(e11, op1(e10, e13))) & ~ (e10 = op1(e10, op1(e10, e13))) & (op1(e10, e13) = op1(e13, e10))) | ~ sP18), inference(nnf_transformation, [], [f39])).
fof(f1653, plain, (~ spl72_190 | ~ spl72_193), inference(avatar_split_clause, [], [f445, f1650, f1636])).
fof(f445, plain, (~ (e11 = op1(e11, op1(e10, e13))) | ~ sP18), inference(cnf_transformation, [], [f107])).
fof(f1634, plain, (~ spl72_184 | spl72_189), inference(avatar_split_clause, [], [f438, f1631, f1607])).
fof(f438, plain, ((op1(e10, e11) = op1(e11, e10)) | ~ sP19), inference(cnf_transformation, [], [f106])).
fof(f106, plain, ((~ (e13 = op1(e13, op1(e11, e10))) & ~ (e12 = op1(e12, op1(e11, e10))) & ~ (e11 = op1(e11, op1(e11, e10))) & ~ (e10 = op1(e10, op1(e11, e10))) & (op1(e10, e11) = op1(e11, e10))) | ~ sP19), inference(nnf_transformation, [], [f40])).
fof(f1600, plain, (~ spl72_179 | ~ spl72_182), inference(avatar_split_clause, [], [f435, f1597, f1583])).
fof(f435, plain, (~ (e11 = op1(e11, op1(e11, e11))) | ~ sP20), inference(cnf_transformation, [], [f105])).
fof(f105, plain, ((~ (e13 = op1(e13, op1(e11, e11))) & ~ (e12 = op1(e12, op1(e11, e11))) & ~ (e11 = op1(e11, op1(e11, e11))) & ~ (e10 = op1(e10, op1(e11, e11))) & (op1(e11, e11) = op1(e11, e11))) | ~ sP20), inference(nnf_transformation, [], [f41])).
fof(f1595, plain, (~ spl72_179 | ~ spl72_181), inference(avatar_split_clause, [], [f436, f1592, f1583])).
fof(f436, plain, (~ (e12 = op1(e12, op1(e11, e11))) | ~ sP20), inference(cnf_transformation, [], [f105])).
fof(f1580, plain, (~ spl72_174 | ~ spl72_178), inference(avatar_split_clause, [], [f429, f1577, f1558])).
fof(f429, plain, (~ (e10 = op1(e10, op1(e11, e12))) | ~ sP21), inference(cnf_transformation, [], [f104])).
fof(f104, plain, ((~ (e13 = op1(e13, op1(e11, e12))) & ~ (e12 = op1(e12, op1(e11, e12))) & ~ (e11 = op1(e11, op1(e11, e12))) & ~ (e10 = op1(e10, op1(e11, e12))) & (op1(e11, e12) = op1(e12, e11))) | ~ sP21), inference(nnf_transformation, [], [f42])).
fof(f1575, plain, (~ spl72_174 | ~ spl72_177), inference(avatar_split_clause, [], [f430, f1572, f1558])).
fof(f430, plain, (~ (e11 = op1(e11, op1(e11, e12))) | ~ sP21), inference(cnf_transformation, [], [f104])).
fof(f1556, plain, (~ spl72_169 | spl72_140), inference(avatar_split_clause, [], [f423, f1392, f1533])).
fof(f423, plain, ((op1(e11, e13) = op1(e13, e11)) | ~ sP22), inference(cnf_transformation, [], [f103])).
fof(f103, plain, ((~ (e13 = op1(e13, op1(e11, e13))) & ~ (e12 = op1(e12, op1(e11, e13))) & ~ (e11 = op1(e11, op1(e11, e13))) & ~ (e10 = op1(e10, op1(e11, e13))) & (op1(e11, e13) = op1(e13, e11))) | ~ sP22), inference(nnf_transformation, [], [f43])).
fof(f1545, plain, (~ spl72_169 | ~ spl72_171), inference(avatar_split_clause, [], [f426, f1542, f1533])).
fof(f426, plain, (~ (e12 = op1(e12, op1(e11, e13))) | ~ sP22), inference(cnf_transformation, [], [f103])).
fof(f1531, plain, (~ spl72_163 | spl72_168), inference(avatar_split_clause, [], [f418, f1528, f1504])).
fof(f418, plain, ((op1(e10, e12) = op1(e12, e10)) | ~ sP23), inference(cnf_transformation, [], [f102])).
fof(f102, plain, ((~ (e13 = op1(e13, op1(e12, e10))) & ~ (e12 = op1(e12, op1(e12, e10))) & ~ (e11 = op1(e11, op1(e12, e10))) & ~ (e10 = op1(e10, op1(e12, e10))) & (op1(e10, e12) = op1(e12, e10))) | ~ sP23), inference(nnf_transformation, [], [f44])).
fof(f1502, plain, (~ spl72_157 | spl72_162), inference(avatar_split_clause, [], [f413, f1499, f1475])).
fof(f413, plain, ((op1(e11, e12) = op1(e12, e11)) | ~ sP24), inference(cnf_transformation, [], [f101])).
fof(f101, plain, ((~ (e13 = op1(e13, op1(e12, e11))) & ~ (e12 = op1(e12, op1(e12, e11))) & ~ (e11 = op1(e11, op1(e12, e11))) & ~ (e10 = op1(e10, op1(e12, e11))) & (op1(e11, e12) = op1(e12, e11))) | ~ sP24), inference(nnf_transformation, [], [f45])).
fof(f1492, plain, (~ spl72_157 | ~ spl72_160), inference(avatar_split_clause, [], [f415, f1489, f1475])).
fof(f415, plain, (~ (e11 = op1(e11, op1(e12, e11))) | ~ sP24), inference(cnf_transformation, [], [f101])).
fof(f1463, plain, (~ spl72_152 | ~ spl72_154), inference(avatar_split_clause, [], [f411, f1460, f1451])).
fof(f411, plain, (~ (e12 = op1(e12, op1(e12, e12))) | ~ sP25), inference(cnf_transformation, [], [f100])).
fof(f100, plain, ((~ (e13 = op1(e13, op1(e12, e12))) & ~ (e12 = op1(e12, op1(e12, e12))) & ~ (e11 = op1(e11, op1(e12, e12))) & ~ (e10 = op1(e10, op1(e12, e12))) & (op1(e12, e12) = op1(e12, e12))) | ~ sP25), inference(nnf_transformation, [], [f46])).
fof(f1449, plain, (~ spl72_147 | spl72_134), inference(avatar_split_clause, [], [f403, f1363, f1426])).
fof(f403, plain, ((op1(e12, e13) = op1(e13, e12)) | ~ sP26), inference(cnf_transformation, [], [f99])).
fof(f99, plain, ((~ (e13 = op1(e13, op1(e12, e13))) & ~ (e12 = op1(e12, op1(e12, e13))) & ~ (e11 = op1(e11, op1(e12, e13))) & ~ (e10 = op1(e10, op1(e12, e13))) & (op1(e12, e13) = op1(e13, e12))) | ~ sP26), inference(nnf_transformation, [], [f47])).
fof(f1424, plain, (~ spl72_141 | spl72_146), inference(avatar_split_clause, [], [f398, f1421, f1397])).
fof(f398, plain, ((op1(e10, e13) = op1(e13, e10)) | ~ sP27), inference(cnf_transformation, [], [f98])).
fof(f98, plain, ((~ (e13 = op1(e13, op1(e13, e10))) & ~ (e12 = op1(e12, op1(e13, e10))) & ~ (e11 = op1(e11, op1(e13, e10))) & ~ (e10 = op1(e10, op1(e13, e10))) & (op1(e10, e13) = op1(e13, e10))) | ~ sP27), inference(nnf_transformation, [], [f48])).
fof(f1414, plain, (~ spl72_141 | ~ spl72_144), inference(avatar_split_clause, [], [f400, f1411, f1397])).
fof(f400, plain, (~ (e11 = op1(e11, op1(e13, e10))) | ~ sP27), inference(cnf_transformation, [], [f98])).
fof(f1395, plain, (~ spl72_135 | spl72_140), inference(avatar_split_clause, [], [f393, f1392, f1368])).
fof(f393, plain, ((op1(e11, e13) = op1(e13, e11)) | ~ sP28), inference(cnf_transformation, [], [f97])).
fof(f97, plain, ((~ (e13 = op1(e13, op1(e13, e11))) & ~ (e12 = op1(e12, op1(e13, e11))) & ~ (e11 = op1(e11, op1(e13, e11))) & ~ (e10 = op1(e10, op1(e13, e11))) & (op1(e11, e13) = op1(e13, e11))) | ~ sP28), inference(nnf_transformation, [], [f49])).
fof(f1380, plain, (~ spl72_135 | ~ spl72_137), inference(avatar_split_clause, [], [f396, f1377, f1368])).
fof(f396, plain, (~ (e12 = op1(e12, op1(e13, e11))) | ~ sP28), inference(cnf_transformation, [], [f97])).
fof(f1366, plain, (~ spl72_129 | spl72_134), inference(avatar_split_clause, [], [f388, f1363, f1339])).
fof(f388, plain, ((op1(e12, e13) = op1(e13, e12)) | ~ sP29), inference(cnf_transformation, [], [f96])).
fof(f96, plain, ((~ (e13 = op1(e13, op1(e13, e12))) & ~ (e12 = op1(e12, op1(e13, e12))) & ~ (e11 = op1(e11, op1(e13, e12))) & ~ (e10 = op1(e10, op1(e13, e12))) & (op1(e12, e13) = op1(e13, e12))) | ~ sP29), inference(nnf_transformation, [], [f50])).
fof(f1330, plain, (spl72_128 | spl72_112 | spl72_96 | spl72_80), inference(avatar_split_clause, [], [f239, f1098, f1166, f1234, f1302])).
fof(f239, plain, ((e23 = op2(e23, e20)) | (e23 = op2(e22, e20)) | (e23 = op2(e21, e20)) | (op2(e20, e20) = e23)), inference(cnf_transformation, [], [f4])).
fof(f4, plain, (((e23 = op2(e23, e23)) | (e23 = op2(e22, e23)) | (e23 = op2(e21, e23)) | (e23 = op2(e20, e23))) & ((e23 = op2(e23, e23)) | (e23 = op2(e23, e22)) | (e23 = op2(e23, e21)) | (e23 = op2(e23, e20))) & ((e22 = op2(e23, e23)) | (e22 = op2(e22, e23)) | (e22 = op2(e21, e23)) | (e22 = op2(e20, e23))) & ((e22 = op2(e23, e23)) | (e22 = op2(e23, e22)) | (e22 = op2(e23, e21)) | (e22 = op2(e23, e20))) & ((e21 = op2(e23, e23)) | (e21 = op2(e22, e23)) | (e21 = op2(e21, e23)) | (e21 = op2(e20, e23))) & ((e21 = op2(e23, e23)) | (e21 = op2(e23, e22)) | (e21 = op2(e23, e21)) | (e21 = op2(e23, e20))) & ((e20 = op2(e23, e23)) | (e20 = op2(e22, e23)) | (e20 = op2(e21, e23)) | (e20 = op2(e20, e23))) & ((e20 = op2(e23, e23)) | (e20 = op2(e23, e22)) | (e20 = op2(e23, e21)) | (e20 = op2(e23, e20))) & ((e23 = op2(e23, e22)) | (e23 = op2(e22, e22)) | (e23 = op2(e21, e22)) | (e23 = op2(e20, e22))) & ((e23 = op2(e22, e23)) | (e23 = op2(e22, e22)) | (e23 = op2(e22, e21)) | (e23 = op2(e22, e20))) & ((e22 = op2(e23, e22)) | (e22 = op2(e22, e22)) | (e22 = op2(e21, e22)) | (e22 = op2(e20, e22))) & ((e22 = op2(e22, e23)) | (e22 = op2(e22, e22)) | (e22 = op2(e22, e21)) | (e22 = op2(e22, e20))) & ((e21 = op2(e23, e22)) | (e21 = op2(e22, e22)) | (e21 = op2(e21, e22)) | (e21 = op2(e20, e22))) & ((e21 = op2(e22, e23)) | (e21 = op2(e22, e22)) | (e21 = op2(e22, e21)) | (e21 = op2(e22, e20))) & ((e20 = op2(e23, e22)) | (e20 = op2(e22, e22)) | (e20 = op2(e21, e22)) | (e20 = op2(e20, e22))) & ((e20 = op2(e22, e23)) | (e20 = op2(e22, e22)) | (e20 = op2(e22, e21)) | (e20 = op2(e22, e20))) & ((e23 = op2(e23, e21)) | (e23 = op2(e22, e21)) | (e23 = op2(e21, e21)) | (e23 = op2(e20, e21))) & ((e23 = op2(e21, e23)) | (e23 = op2(e21, e22)) | (e23 = op2(e21, e21)) | (e23 = op2(e21, e20))) & ((e22 = op2(e23, e21)) | (e22 = op2(e22, e21)) | (e22 = op2(e21, e21)) | (e22 = op2(e20, e21))) & ((e22 = op2(e21, e23)) | (e22 = op2(e21, e22)) | (e22 = op2(e21, e21)) | (e22 = op2(e21, e20))) & ((e21 = op2(e23, e21)) | (e21 = op2(e22, e21)) | (e21 = op2(e21, e21)) | (e21 = op2(e20, e21))) & ((e21 = op2(e21, e23)) | (e21 = op2(e21, e22)) | (e21 = op2(e21, e21)) | (e21 = op2(e21, e20))) & ((e20 = op2(e23, e21)) | (e20 = op2(e22, e21)) | (e20 = op2(e21, e21)) | (e20 = op2(e20, e21))) & ((e20 = op2(e21, e23)) | (e20 = op2(e21, e22)) | (e20 = op2(e21, e21)) | (e20 = op2(e21, e20))) & ((e23 = op2(e23, e20)) | (e23 = op2(e22, e20)) | (e23 = op2(e21, e20)) | (op2(e20, e20) = e23)) & ((e23 = op2(e20, e23)) | (e23 = op2(e20, e22)) | (e23 = op2(e20, e21)) | (op2(e20, e20) = e23)) & ((e22 = op2(e23, e20)) | (e22 = op2(e22, e20)) | (e22 = op2(e21, e20)) | (op2(e20, e20) = e22)) & ((e22 = op2(e20, e23)) | (e22 = op2(e20, e22)) | (e22 = op2(e20, e21)) | (op2(e20, e20) = e22)) & ((e21 = op2(e23, e20)) | (e21 = op2(e22, e20)) | (e21 = op2(e21, e20)) | (op2(e20, e20) = e21)) & ((e21 = op2(e20, e23)) | (e21 = op2(e20, e22)) | (e21 = op2(e20, e21)) | (op2(e20, e20) = e21)) & ((e20 = op2(e23, e20)) | (e20 = op2(e22, e20)) | (e20 = op2(e21, e20)) | (e20 = op2(e20, e20))) & ((e20 = op2(e20, e23)) | (e20 = op2(e20, e22)) | (e20 = op2(e20, e21)) | (e20 = op2(e20, e20)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG104+1.p', ax4)).
fof(f1329, plain, (spl72_109 | spl72_105 | spl72_101 | spl72_97), inference(avatar_split_clause, [], [f240, f1171, f1188, f1205, f1222])).
fof(f240, plain, ((e20 = op2(e21, e23)) | (e20 = op2(e21, e22)) | (e20 = op2(e21, e21)) | (e20 = op2(e21, e20))), inference(cnf_transformation, [], [f4])).
fof(f1328, plain, (spl72_121 | spl72_105 | spl72_89 | spl72_73), inference(avatar_split_clause, [], [f241, f1069, f1137, f1205, f1273])).
fof(f241, plain, ((e20 = op2(e23, e21)) | (e20 = op2(e22, e21)) | (e20 = op2(e21, e21)) | (e20 = op2(e20, e21))), inference(cnf_transformation, [], [f4])).
fof(f1327, plain, (spl72_110 | spl72_106 | spl72_102 | spl72_98), inference(avatar_split_clause, [], [f242, f1175, f1192, f1209, f1226])).
fof(f242, plain, ((e21 = op2(e21, e23)) | (e21 = op2(e21, e22)) | (e21 = op2(e21, e21)) | (e21 = op2(e21, e20))), inference(cnf_transformation, [], [f4])).
fof(f1326, plain, (spl72_122 | spl72_106 | spl72_90 | spl72_74), inference(avatar_split_clause, [], [f243, f1073, f1141, f1209, f1277])).
fof(f243, plain, ((e21 = op2(e23, e21)) | (e21 = op2(e22, e21)) | (e21 = op2(e21, e21)) | (e21 = op2(e20, e21))), inference(cnf_transformation, [], [f4])).
fof(f1325, plain, (spl72_111 | spl72_107 | spl72_103 | spl72_99), inference(avatar_split_clause, [], [f244, f1179, f1196, f1213, f1230])).
fof(f244, plain, ((e22 = op2(e21, e23)) | (e22 = op2(e21, e22)) | (e22 = op2(e21, e21)) | (e22 = op2(e21, e20))), inference(cnf_transformation, [], [f4])).
fof(f1323, plain, (spl72_112 | spl72_108 | spl72_104 | spl72_100), inference(avatar_split_clause, [], [f246, f1183, f1200, f1217, f1234])).
fof(f246, plain, ((e23 = op2(e21, e23)) | (e23 = op2(e21, e22)) | (e23 = op2(e21, e21)) | (e23 = op2(e21, e20))), inference(cnf_transformation, [], [f4])).
fof(f1321, plain, (spl72_93 | spl72_89 | spl72_85 | spl72_81), inference(avatar_split_clause, [], [f248, f1103, f1120, f1137, f1154])).
fof(f248, plain, ((e20 = op2(e22, e23)) | (e20 = op2(e22, e22)) | (e20 = op2(e22, e21)) | (e20 = op2(e22, e20))), inference(cnf_transformation, [], [f4])).
fof(f1320, plain, (spl72_117 | spl72_101 | spl72_85 | spl72_69), inference(avatar_split_clause, [], [f249, f1052, f1120, f1188, f1256])).
fof(f249, plain, ((e20 = op2(e23, e22)) | (e20 = op2(e22, e22)) | (e20 = op2(e21, e22)) | (e20 = op2(e20, e22))), inference(cnf_transformation, [], [f4])).
fof(f1319, plain, (spl72_94 | spl72_90 | spl72_86 | spl72_82), inference(avatar_split_clause, [], [f250, f1107, f1124, f1141, f1158])).
fof(f250, plain, ((e21 = op2(e22, e23)) | (e21 = op2(e22, e22)) | (e21 = op2(e22, e21)) | (e21 = op2(e22, e20))), inference(cnf_transformation, [], [f4])).
fof(f1317, plain, (spl72_95 | spl72_91 | spl72_87 | spl72_83), inference(avatar_split_clause, [], [f252, f1111, f1128, f1145, f1162])).
fof(f252, plain, ((e22 = op2(e22, e23)) | (e22 = op2(e22, e22)) | (e22 = op2(e22, e21)) | (e22 = op2(e22, e20))), inference(cnf_transformation, [], [f4])).
fof(f1316, plain, (spl72_119 | spl72_103 | spl72_87 | spl72_71), inference(avatar_split_clause, [], [f253, f1060, f1128, f1196, f1264])).
fof(f253, plain, ((e22 = op2(e23, e22)) | (e22 = op2(e22, e22)) | (e22 = op2(e21, e22)) | (e22 = op2(e20, e22))), inference(cnf_transformation, [], [f4])).
fof(f1315, plain, (spl72_96 | spl72_92 | spl72_88 | spl72_84), inference(avatar_split_clause, [], [f254, f1115, f1132, f1149, f1166])).
fof(f254, plain, ((e23 = op2(e22, e23)) | (e23 = op2(e22, e22)) | (e23 = op2(e22, e21)) | (e23 = op2(e22, e20))), inference(cnf_transformation, [], [f4])).
fof(f1314, plain, (spl72_120 | spl72_104 | spl72_88 | spl72_72), inference(avatar_split_clause, [], [f255, f1064, f1132, f1200, f1268])).
fof(f255, plain, ((e23 = op2(e23, e22)) | (e23 = op2(e22, e22)) | (e23 = op2(e21, e22)) | (e23 = op2(e20, e22))), inference(cnf_transformation, [], [f4])).
fof(f1313, plain, (spl72_77 | spl72_73 | spl72_69 | spl72_65), inference(avatar_split_clause, [], [f256, f1035, f1052, f1069, f1086])).
fof(f256, plain, ((e20 = op2(e23, e23)) | (e20 = op2(e23, e22)) | (e20 = op2(e23, e21)) | (e20 = op2(e23, e20))), inference(cnf_transformation, [], [f4])).
fof(f1311, plain, (spl72_78 | spl72_74 | spl72_70 | spl72_66), inference(avatar_split_clause, [], [f258, f1039, f1056, f1073, f1090])).
fof(f258, plain, ((e21 = op2(e23, e23)) | (e21 = op2(e23, e22)) | (e21 = op2(e23, e21)) | (e21 = op2(e23, e20))), inference(cnf_transformation, [], [f4])).
fof(f1310, plain, (spl72_114 | spl72_98 | spl72_82 | spl72_66), inference(avatar_split_clause, [], [f259, f1039, f1107, f1175, f1243])).
fof(f259, plain, ((e21 = op2(e23, e23)) | (e21 = op2(e22, e23)) | (e21 = op2(e21, e23)) | (e21 = op2(e20, e23))), inference(cnf_transformation, [], [f4])).
fof(f1309, plain, (spl72_79 | spl72_75 | spl72_71 | spl72_67), inference(avatar_split_clause, [], [f260, f1043, f1060, f1077, f1094])).
fof(f260, plain, ((e22 = op2(e23, e23)) | (e22 = op2(e23, e22)) | (e22 = op2(e23, e21)) | (e22 = op2(e23, e20))), inference(cnf_transformation, [], [f4])).
fof(f1308, plain, (spl72_115 | spl72_99 | spl72_83 | spl72_67), inference(avatar_split_clause, [], [f261, f1043, f1111, f1179, f1247])).
fof(f261, plain, ((e22 = op2(e23, e23)) | (e22 = op2(e22, e23)) | (e22 = op2(e21, e23)) | (e22 = op2(e20, e23))), inference(cnf_transformation, [], [f4])).
fof(f1307, plain, (spl72_80 | spl72_76 | spl72_72 | spl72_68), inference(avatar_split_clause, [], [f262, f1047, f1064, f1081, f1098])).
fof(f262, plain, ((e23 = op2(e23, e23)) | (e23 = op2(e23, e22)) | (e23 = op2(e23, e21)) | (e23 = op2(e23, e20))), inference(cnf_transformation, [], [f4])).
fof(f1306, plain, (spl72_116 | spl72_100 | spl72_84 | spl72_68), inference(avatar_split_clause, [], [f263, f1047, f1115, f1183, f1251])).
fof(f263, plain, ((e23 = op2(e23, e23)) | (e23 = op2(e22, e23)) | (e23 = op2(e21, e23)) | (e23 = op2(e20, e23))), inference(cnf_transformation, [], [f4])).
fof(f1254, plain, (spl72_113 | spl72_114 | spl72_115 | spl72_116), inference(avatar_split_clause, [], [f219, f1251, f1247, f1243, f1239])).
fof(f219, plain, ((e23 = op2(e20, e23)) | (e22 = op2(e20, e23)) | (e21 = op2(e20, e23)) | (e20 = op2(e20, e23))), inference(cnf_transformation, [], [f3])).
fof(f3, plain, (((e23 = op2(e23, e23)) | (e22 = op2(e23, e23)) | (e21 = op2(e23, e23)) | (e20 = op2(e23, e23))) & ((e23 = op2(e23, e22)) | (e22 = op2(e23, e22)) | (e21 = op2(e23, e22)) | (e20 = op2(e23, e22))) & ((e23 = op2(e23, e21)) | (e22 = op2(e23, e21)) | (e21 = op2(e23, e21)) | (e20 = op2(e23, e21))) & ((e23 = op2(e23, e20)) | (e22 = op2(e23, e20)) | (e21 = op2(e23, e20)) | (e20 = op2(e23, e20))) & ((e23 = op2(e22, e23)) | (e22 = op2(e22, e23)) | (e21 = op2(e22, e23)) | (e20 = op2(e22, e23))) & ((e23 = op2(e22, e22)) | (e22 = op2(e22, e22)) | (e21 = op2(e22, e22)) | (e20 = op2(e22, e22))) & ((e23 = op2(e22, e21)) | (e22 = op2(e22, e21)) | (e21 = op2(e22, e21)) | (e20 = op2(e22, e21))) & ((e23 = op2(e22, e20)) | (e22 = op2(e22, e20)) | (e21 = op2(e22, e20)) | (e20 = op2(e22, e20))) & ((e23 = op2(e21, e23)) | (e22 = op2(e21, e23)) | (e21 = op2(e21, e23)) | (e20 = op2(e21, e23))) & ((e23 = op2(e21, e22)) | (e22 = op2(e21, e22)) | (e21 = op2(e21, e22)) | (e20 = op2(e21, e22))) & ((e23 = op2(e21, e21)) | (e22 = op2(e21, e21)) | (e21 = op2(e21, e21)) | (e20 = op2(e21, e21))) & ((e23 = op2(e21, e20)) | (e22 = op2(e21, e20)) | (e21 = op2(e21, e20)) | (e20 = op2(e21, e20))) & ((e23 = op2(e20, e23)) | (e22 = op2(e20, e23)) | (e21 = op2(e20, e23)) | (e20 = op2(e20, e23))) & ((e23 = op2(e20, e22)) | (e22 = op2(e20, e22)) | (e21 = op2(e20, e22)) | (e20 = op2(e20, e22))) & ((e23 = op2(e20, e21)) | (e22 = op2(e20, e21)) | (e21 = op2(e20, e21)) | (e20 = op2(e20, e21))) & ((op2(e20, e20) = e23) | (op2(e20, e20) = e22) | (op2(e20, e20) = e21) | (e20 = op2(e20, e20)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG104+1.p', ax3)).
fof(f1237, plain, (spl72_109 | spl72_110 | spl72_111 | spl72_112), inference(avatar_split_clause, [], [f220, f1234, f1230, f1226, f1222])).
fof(f220, plain, ((e23 = op2(e21, e20)) | (e22 = op2(e21, e20)) | (e21 = op2(e21, e20)) | (e20 = op2(e21, e20))), inference(cnf_transformation, [], [f3])).
fof(f1220, plain, (spl72_105 | spl72_106 | spl72_107 | spl72_108), inference(avatar_split_clause, [], [f221, f1217, f1213, f1209, f1205])).
fof(f221, plain, ((e23 = op2(e21, e21)) | (e22 = op2(e21, e21)) | (e21 = op2(e21, e21)) | (e20 = op2(e21, e21))), inference(cnf_transformation, [], [f3])).
fof(f1203, plain, (spl72_101 | spl72_102 | spl72_103 | spl72_104), inference(avatar_split_clause, [], [f222, f1200, f1196, f1192, f1188])).
fof(f222, plain, ((e23 = op2(e21, e22)) | (e22 = op2(e21, e22)) | (e21 = op2(e21, e22)) | (e20 = op2(e21, e22))), inference(cnf_transformation, [], [f3])).
fof(f1186, plain, (spl72_97 | spl72_98 | spl72_99 | spl72_100), inference(avatar_split_clause, [], [f223, f1183, f1179, f1175, f1171])).
fof(f223, plain, ((e23 = op2(e21, e23)) | (e22 = op2(e21, e23)) | (e21 = op2(e21, e23)) | (e20 = op2(e21, e23))), inference(cnf_transformation, [], [f3])).
fof(f1169, plain, (spl72_93 | spl72_94 | spl72_95 | spl72_96), inference(avatar_split_clause, [], [f224, f1166, f1162, f1158, f1154])).
fof(f224, plain, ((e23 = op2(e22, e20)) | (e22 = op2(e22, e20)) | (e21 = op2(e22, e20)) | (e20 = op2(e22, e20))), inference(cnf_transformation, [], [f3])).
fof(f1152, plain, (spl72_89 | spl72_90 | spl72_91 | spl72_92), inference(avatar_split_clause, [], [f225, f1149, f1145, f1141, f1137])).
fof(f225, plain, ((e23 = op2(e22, e21)) | (e22 = op2(e22, e21)) | (e21 = op2(e22, e21)) | (e20 = op2(e22, e21))), inference(cnf_transformation, [], [f3])).
fof(f1135, plain, (spl72_85 | spl72_86 | spl72_87 | spl72_88), inference(avatar_split_clause, [], [f226, f1132, f1128, f1124, f1120])).
fof(f226, plain, ((e23 = op2(e22, e22)) | (e22 = op2(e22, e22)) | (e21 = op2(e22, e22)) | (e20 = op2(e22, e22))), inference(cnf_transformation, [], [f3])).
fof(f1118, plain, (spl72_81 | spl72_82 | spl72_83 | spl72_84), inference(avatar_split_clause, [], [f227, f1115, f1111, f1107, f1103])).
fof(f227, plain, ((e23 = op2(e22, e23)) | (e22 = op2(e22, e23)) | (e21 = op2(e22, e23)) | (e20 = op2(e22, e23))), inference(cnf_transformation, [], [f3])).
fof(f1101, plain, (spl72_77 | spl72_78 | spl72_79 | spl72_80), inference(avatar_split_clause, [], [f228, f1098, f1094, f1090, f1086])).
fof(f228, plain, ((e23 = op2(e23, e20)) | (e22 = op2(e23, e20)) | (e21 = op2(e23, e20)) | (e20 = op2(e23, e20))), inference(cnf_transformation, [], [f3])).
fof(f1084, plain, (spl72_73 | spl72_74 | spl72_75 | spl72_76), inference(avatar_split_clause, [], [f229, f1081, f1077, f1073, f1069])).
fof(f229, plain, ((e23 = op2(e23, e21)) | (e22 = op2(e23, e21)) | (e21 = op2(e23, e21)) | (e20 = op2(e23, e21))), inference(cnf_transformation, [], [f3])).
fof(f1067, plain, (spl72_69 | spl72_70 | spl72_71 | spl72_72), inference(avatar_split_clause, [], [f230, f1064, f1060, f1056, f1052])).
fof(f230, plain, ((e23 = op2(e23, e22)) | (e22 = op2(e23, e22)) | (e21 = op2(e23, e22)) | (e20 = op2(e23, e22))), inference(cnf_transformation, [], [f3])).
fof(f1050, plain, (spl72_65 | spl72_66 | spl72_67 | spl72_68), inference(avatar_split_clause, [], [f231, f1047, f1043, f1039, f1035])).
fof(f231, plain, ((e23 = op2(e23, e23)) | (e22 = op2(e23, e23)) | (e21 = op2(e23, e23)) | (e20 = op2(e23, e23))), inference(cnf_transformation, [], [f3])).
fof(f1026, plain, (spl72_64 | spl72_48 | spl72_32 | spl72_16), inference(avatar_split_clause, [], [f191, f794, f862, f930, f998])).
fof(f191, plain, ((e13 = op1(e13, e10)) | (e13 = op1(e12, e10)) | (e13 = op1(e11, e10)) | (op1(e10, e10) = e13)), inference(cnf_transformation, [], [f2])).
fof(f2, plain, (((e13 = op1(e13, e13)) | (e13 = op1(e12, e13)) | (e13 = op1(e11, e13)) | (e13 = op1(e10, e13))) & ((e13 = op1(e13, e13)) | (e13 = op1(e13, e12)) | (e13 = op1(e13, e11)) | (e13 = op1(e13, e10))) & ((e12 = op1(e13, e13)) | (e12 = op1(e12, e13)) | (e12 = op1(e11, e13)) | (e12 = op1(e10, e13))) & ((e12 = op1(e13, e13)) | (e12 = op1(e13, e12)) | (e12 = op1(e13, e11)) | (e12 = op1(e13, e10))) & ((e11 = op1(e13, e13)) | (e11 = op1(e12, e13)) | (e11 = op1(e11, e13)) | (e11 = op1(e10, e13))) & ((e11 = op1(e13, e13)) | (e11 = op1(e13, e12)) | (e11 = op1(e13, e11)) | (e11 = op1(e13, e10))) & ((e10 = op1(e13, e13)) | (e10 = op1(e12, e13)) | (e10 = op1(e11, e13)) | (e10 = op1(e10, e13))) & ((e10 = op1(e13, e13)) | (e10 = op1(e13, e12)) | (e10 = op1(e13, e11)) | (e10 = op1(e13, e10))) & ((e13 = op1(e13, e12)) | (e13 = op1(e12, e12)) | (e13 = op1(e11, e12)) | (e13 = op1(e10, e12))) & ((e13 = op1(e12, e13)) | (e13 = op1(e12, e12)) | (e13 = op1(e12, e11)) | (e13 = op1(e12, e10))) & ((e12 = op1(e13, e12)) | (e12 = op1(e12, e12)) | (e12 = op1(e11, e12)) | (e12 = op1(e10, e12))) & ((e12 = op1(e12, e13)) | (e12 = op1(e12, e12)) | (e12 = op1(e12, e11)) | (e12 = op1(e12, e10))) & ((e11 = op1(e13, e12)) | (e11 = op1(e12, e12)) | (e11 = op1(e11, e12)) | (e11 = op1(e10, e12))) & ((e11 = op1(e12, e13)) | (e11 = op1(e12, e12)) | (e11 = op1(e12, e11)) | (e11 = op1(e12, e10))) & ((e10 = op1(e13, e12)) | (e10 = op1(e12, e12)) | (e10 = op1(e11, e12)) | (e10 = op1(e10, e12))) & ((e10 = op1(e12, e13)) | (e10 = op1(e12, e12)) | (e10 = op1(e12, e11)) | (e10 = op1(e12, e10))) & ((e13 = op1(e13, e11)) | (e13 = op1(e12, e11)) | (e13 = op1(e11, e11)) | (e13 = op1(e10, e11))) & ((e13 = op1(e11, e13)) | (e13 = op1(e11, e12)) | (e13 = op1(e11, e11)) | (e13 = op1(e11, e10))) & ((e12 = op1(e13, e11)) | (e12 = op1(e12, e11)) | (e12 = op1(e11, e11)) | (e12 = op1(e10, e11))) & ((e12 = op1(e11, e13)) | (e12 = op1(e11, e12)) | (e12 = op1(e11, e11)) | (e12 = op1(e11, e10))) & ((e11 = op1(e13, e11)) | (e11 = op1(e12, e11)) | (e11 = op1(e11, e11)) | (e11 = op1(e10, e11))) & ((e11 = op1(e11, e13)) | (e11 = op1(e11, e12)) | (e11 = op1(e11, e11)) | (e11 = op1(e11, e10))) & ((e10 = op1(e13, e11)) | (e10 = op1(e12, e11)) | (e10 = op1(e11, e11)) | (e10 = op1(e10, e11))) & ((e10 = op1(e11, e13)) | (e10 = op1(e11, e12)) | (e10 = op1(e11, e11)) | (e10 = op1(e11, e10))) & ((e13 = op1(e13, e10)) | (e13 = op1(e12, e10)) | (e13 = op1(e11, e10)) | (op1(e10, e10) = e13)) & ((e13 = op1(e10, e13)) | (e13 = op1(e10, e12)) | (e13 = op1(e10, e11)) | (op1(e10, e10) = e13)) & ((e12 = op1(e13, e10)) | (e12 = op1(e12, e10)) | (e12 = op1(e11, e10)) | (op1(e10, e10) = e12)) & ((e12 = op1(e10, e13)) | (e12 = op1(e10, e12)) | (e12 = op1(e10, e11)) | (op1(e10, e10) = e12)) & ((e11 = op1(e13, e10)) | (e11 = op1(e12, e10)) | (e11 = op1(e11, e10)) | (op1(e10, e10) = e11)) & ((e11 = op1(e10, e13)) | (e11 = op1(e10, e12)) | (e11 = op1(e10, e11)) | (op1(e10, e10) = e11)) & ((e10 = op1(e13, e10)) | (e10 = op1(e12, e10)) | (e10 = op1(e11, e10)) | (e10 = op1(e10, e10))) & ((e10 = op1(e10, e13)) | (e10 = op1(e10, e12)) | (e10 = op1(e10, e11)) | (e10 = op1(e10, e10)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG104+1.p', ax2)).
fof(f1024, plain, (spl72_57 | spl72_41 | spl72_25 | spl72_9), inference(avatar_split_clause, [], [f193, f765, f833, f901, f969])).
fof(f193, plain, ((e10 = op1(e13, e11)) | (e10 = op1(e12, e11)) | (e10 = op1(e11, e11)) | (e10 = op1(e10, e11))), inference(cnf_transformation, [], [f2])).
fof(f1023, plain, (spl72_46 | spl72_42 | spl72_38 | spl72_34), inference(avatar_split_clause, [], [f194, f871, f888, f905, f922])).
fof(f194, plain, ((e11 = op1(e11, e13)) | (e11 = op1(e11, e12)) | (e11 = op1(e11, e11)) | (e11 = op1(e11, e10))), inference(cnf_transformation, [], [f2])).
fof(f1022, plain, (spl72_58 | spl72_42 | spl72_26 | spl72_10), inference(avatar_split_clause, [], [f195, f769, f837, f905, f973])).
fof(f195, plain, ((e11 = op1(e13, e11)) | (e11 = op1(e12, e11)) | (e11 = op1(e11, e11)) | (e11 = op1(e10, e11))), inference(cnf_transformation, [], [f2])).
fof(f1021, plain, (spl72_47 | spl72_43 | spl72_39 | spl72_35), inference(avatar_split_clause, [], [f196, f875, f892, f909, f926])).
fof(f196, plain, ((e12 = op1(e11, e13)) | (e12 = op1(e11, e12)) | (e12 = op1(e11, e11)) | (e12 = op1(e11, e10))), inference(cnf_transformation, [], [f2])).
fof(f1020, plain, (spl72_59 | spl72_43 | spl72_27 | spl72_11), inference(avatar_split_clause, [], [f197, f773, f841, f909, f977])).
fof(f197, plain, ((e12 = op1(e13, e11)) | (e12 = op1(e12, e11)) | (e12 = op1(e11, e11)) | (e12 = op1(e10, e11))), inference(cnf_transformation, [], [f2])).
fof(f1019, plain, (spl72_48 | spl72_44 | spl72_40 | spl72_36), inference(avatar_split_clause, [], [f198, f879, f896, f913, f930])).
fof(f198, plain, ((e13 = op1(e11, e13)) | (e13 = op1(e11, e12)) | (e13 = op1(e11, e11)) | (e13 = op1(e11, e10))), inference(cnf_transformation, [], [f2])).
fof(f1017, plain, (spl72_29 | spl72_25 | spl72_21 | spl72_17), inference(avatar_split_clause, [], [f200, f799, f816, f833, f850])).
fof(f200, plain, ((e10 = op1(e12, e13)) | (e10 = op1(e12, e12)) | (e10 = op1(e12, e11)) | (e10 = op1(e12, e10))), inference(cnf_transformation, [], [f2])).
fof(f1016, plain, (spl72_53 | spl72_37 | spl72_21 | spl72_5), inference(avatar_split_clause, [], [f201, f748, f816, f884, f952])).
fof(f201, plain, ((e10 = op1(e13, e12)) | (e10 = op1(e12, e12)) | (e10 = op1(e11, e12)) | (e10 = op1(e10, e12))), inference(cnf_transformation, [], [f2])).
fof(f1015, plain, (spl72_30 | spl72_26 | spl72_22 | spl72_18), inference(avatar_split_clause, [], [f202, f803, f820, f837, f854])).
fof(f202, plain, ((e11 = op1(e12, e13)) | (e11 = op1(e12, e12)) | (e11 = op1(e12, e11)) | (e11 = op1(e12, e10))), inference(cnf_transformation, [], [f2])).
fof(f1013, plain, (spl72_31 | spl72_27 | spl72_23 | spl72_19), inference(avatar_split_clause, [], [f204, f807, f824, f841, f858])).
fof(f204, plain, ((e12 = op1(e12, e13)) | (e12 = op1(e12, e12)) | (e12 = op1(e12, e11)) | (e12 = op1(e12, e10))), inference(cnf_transformation, [], [f2])).
fof(f1012, plain, (spl72_55 | spl72_39 | spl72_23 | spl72_7), inference(avatar_split_clause, [], [f205, f756, f824, f892, f960])).
fof(f205, plain, ((e12 = op1(e13, e12)) | (e12 = op1(e12, e12)) | (e12 = op1(e11, e12)) | (e12 = op1(e10, e12))), inference(cnf_transformation, [], [f2])).
fof(f1010, plain, (spl72_56 | spl72_40 | spl72_24 | spl72_8), inference(avatar_split_clause, [], [f207, f760, f828, f896, f964])).
fof(f207, plain, ((e13 = op1(e13, e12)) | (e13 = op1(e12, e12)) | (e13 = op1(e11, e12)) | (e13 = op1(e10, e12))), inference(cnf_transformation, [], [f2])).
fof(f1009, plain, (spl72_13 | spl72_9 | spl72_5 | spl72_1), inference(avatar_split_clause, [], [f208, f731, f748, f765, f782])).
fof(f208, plain, ((e10 = op1(e13, e13)) | (e10 = op1(e13, e12)) | (e10 = op1(e13, e11)) | (e10 = op1(e13, e10))), inference(cnf_transformation, [], [f2])).
fof(f1007, plain, (spl72_14 | spl72_10 | spl72_6 | spl72_2), inference(avatar_split_clause, [], [f210, f735, f752, f769, f786])).
fof(f210, plain, ((e11 = op1(e13, e13)) | (e11 = op1(e13, e12)) | (e11 = op1(e13, e11)) | (e11 = op1(e13, e10))), inference(cnf_transformation, [], [f2])).
fof(f1006, plain, (spl72_50 | spl72_34 | spl72_18 | spl72_2), inference(avatar_split_clause, [], [f211, f735, f803, f871, f939])).
fof(f211, plain, ((e11 = op1(e13, e13)) | (e11 = op1(e12, e13)) | (e11 = op1(e11, e13)) | (e11 = op1(e10, e13))), inference(cnf_transformation, [], [f2])).
fof(f1005, plain, (spl72_15 | spl72_11 | spl72_7 | spl72_3), inference(avatar_split_clause, [], [f212, f739, f756, f773, f790])).
fof(f212, plain, ((e12 = op1(e13, e13)) | (e12 = op1(e13, e12)) | (e12 = op1(e13, e11)) | (e12 = op1(e13, e10))), inference(cnf_transformation, [], [f2])).
fof(f1004, plain, (spl72_51 | spl72_35 | spl72_19 | spl72_3), inference(avatar_split_clause, [], [f213, f739, f807, f875, f943])).
fof(f213, plain, ((e12 = op1(e13, e13)) | (e12 = op1(e12, e13)) | (e12 = op1(e11, e13)) | (e12 = op1(e10, e13))), inference(cnf_transformation, [], [f2])).
fof(f1003, plain, (spl72_16 | spl72_12 | spl72_8 | spl72_4), inference(avatar_split_clause, [], [f214, f743, f760, f777, f794])).
fof(f214, plain, ((e13 = op1(e13, e13)) | (e13 = op1(e13, e12)) | (e13 = op1(e13, e11)) | (e13 = op1(e13, e10))), inference(cnf_transformation, [], [f2])).
fof(f1002, plain, (spl72_52 | spl72_36 | spl72_20 | spl72_4), inference(avatar_split_clause, [], [f215, f743, f811, f879, f947])).
fof(f215, plain, ((e13 = op1(e13, e13)) | (e13 = op1(e12, e13)) | (e13 = op1(e11, e13)) | (e13 = op1(e10, e13))), inference(cnf_transformation, [], [f2])).
fof(f950, plain, (spl72_49 | spl72_50 | spl72_51 | spl72_52), inference(avatar_split_clause, [], [f171, f947, f943, f939, f935])).
fof(f171, plain, ((e13 = op1(e10, e13)) | (e12 = op1(e10, e13)) | (e11 = op1(e10, e13)) | (e10 = op1(e10, e13))), inference(cnf_transformation, [], [f1])).
fof(f1, plain, (((e13 = op1(e13, e13)) | (e12 = op1(e13, e13)) | (e11 = op1(e13, e13)) | (e10 = op1(e13, e13))) & ((e13 = op1(e13, e12)) | (e12 = op1(e13, e12)) | (e11 = op1(e13, e12)) | (e10 = op1(e13, e12))) & ((e13 = op1(e13, e11)) | (e12 = op1(e13, e11)) | (e11 = op1(e13, e11)) | (e10 = op1(e13, e11))) & ((e13 = op1(e13, e10)) | (e12 = op1(e13, e10)) | (e11 = op1(e13, e10)) | (e10 = op1(e13, e10))) & ((e13 = op1(e12, e13)) | (e12 = op1(e12, e13)) | (e11 = op1(e12, e13)) | (e10 = op1(e12, e13))) & ((e13 = op1(e12, e12)) | (e12 = op1(e12, e12)) | (e11 = op1(e12, e12)) | (e10 = op1(e12, e12))) & ((e13 = op1(e12, e11)) | (e12 = op1(e12, e11)) | (e11 = op1(e12, e11)) | (e10 = op1(e12, e11))) & ((e13 = op1(e12, e10)) | (e12 = op1(e12, e10)) | (e11 = op1(e12, e10)) | (e10 = op1(e12, e10))) & ((e13 = op1(e11, e13)) | (e12 = op1(e11, e13)) | (e11 = op1(e11, e13)) | (e10 = op1(e11, e13))) & ((e13 = op1(e11, e12)) | (e12 = op1(e11, e12)) | (e11 = op1(e11, e12)) | (e10 = op1(e11, e12))) & ((e13 = op1(e11, e11)) | (e12 = op1(e11, e11)) | (e11 = op1(e11, e11)) | (e10 = op1(e11, e11))) & ((e13 = op1(e11, e10)) | (e12 = op1(e11, e10)) | (e11 = op1(e11, e10)) | (e10 = op1(e11, e10))) & ((e13 = op1(e10, e13)) | (e12 = op1(e10, e13)) | (e11 = op1(e10, e13)) | (e10 = op1(e10, e13))) & ((e13 = op1(e10, e12)) | (e12 = op1(e10, e12)) | (e11 = op1(e10, e12)) | (e10 = op1(e10, e12))) & ((e13 = op1(e10, e11)) | (e12 = op1(e10, e11)) | (e11 = op1(e10, e11)) | (e10 = op1(e10, e11))) & ((op1(e10, e10) = e13) | (op1(e10, e10) = e12) | (op1(e10, e10) = e11) | (e10 = op1(e10, e10)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG104+1.p', ax1)).
fof(f933, plain, (spl72_45 | spl72_46 | spl72_47 | spl72_48), inference(avatar_split_clause, [], [f172, f930, f926, f922, f918])).
fof(f172, plain, ((e13 = op1(e11, e10)) | (e12 = op1(e11, e10)) | (e11 = op1(e11, e10)) | (e10 = op1(e11, e10))), inference(cnf_transformation, [], [f1])).
fof(f916, plain, (spl72_41 | spl72_42 | spl72_43 | spl72_44), inference(avatar_split_clause, [], [f173, f913, f909, f905, f901])).
fof(f173, plain, ((e13 = op1(e11, e11)) | (e12 = op1(e11, e11)) | (e11 = op1(e11, e11)) | (e10 = op1(e11, e11))), inference(cnf_transformation, [], [f1])).
fof(f899, plain, (spl72_37 | spl72_38 | spl72_39 | spl72_40), inference(avatar_split_clause, [], [f174, f896, f892, f888, f884])).
fof(f174, plain, ((e13 = op1(e11, e12)) | (e12 = op1(e11, e12)) | (e11 = op1(e11, e12)) | (e10 = op1(e11, e12))), inference(cnf_transformation, [], [f1])).
fof(f882, plain, (spl72_33 | spl72_34 | spl72_35 | spl72_36), inference(avatar_split_clause, [], [f175, f879, f875, f871, f867])).
fof(f175, plain, ((e13 = op1(e11, e13)) | (e12 = op1(e11, e13)) | (e11 = op1(e11, e13)) | (e10 = op1(e11, e13))), inference(cnf_transformation, [], [f1])).
fof(f865, plain, (spl72_29 | spl72_30 | spl72_31 | spl72_32), inference(avatar_split_clause, [], [f176, f862, f858, f854, f850])).
fof(f176, plain, ((e13 = op1(e12, e10)) | (e12 = op1(e12, e10)) | (e11 = op1(e12, e10)) | (e10 = op1(e12, e10))), inference(cnf_transformation, [], [f1])).
fof(f848, plain, (spl72_25 | spl72_26 | spl72_27 | spl72_28), inference(avatar_split_clause, [], [f177, f845, f841, f837, f833])).
fof(f177, plain, ((e13 = op1(e12, e11)) | (e12 = op1(e12, e11)) | (e11 = op1(e12, e11)) | (e10 = op1(e12, e11))), inference(cnf_transformation, [], [f1])).
fof(f831, plain, (spl72_21 | spl72_22 | spl72_23 | spl72_24), inference(avatar_split_clause, [], [f178, f828, f824, f820, f816])).
fof(f178, plain, ((e13 = op1(e12, e12)) | (e12 = op1(e12, e12)) | (e11 = op1(e12, e12)) | (e10 = op1(e12, e12))), inference(cnf_transformation, [], [f1])).
fof(f814, plain, (spl72_17 | spl72_18 | spl72_19 | spl72_20), inference(avatar_split_clause, [], [f179, f811, f807, f803, f799])).
fof(f179, plain, ((e13 = op1(e12, e13)) | (e12 = op1(e12, e13)) | (e11 = op1(e12, e13)) | (e10 = op1(e12, e13))), inference(cnf_transformation, [], [f1])).
fof(f797, plain, (spl72_13 | spl72_14 | spl72_15 | spl72_16), inference(avatar_split_clause, [], [f180, f794, f790, f786, f782])).
fof(f180, plain, ((e13 = op1(e13, e10)) | (e12 = op1(e13, e10)) | (e11 = op1(e13, e10)) | (e10 = op1(e13, e10))), inference(cnf_transformation, [], [f1])).
fof(f780, plain, (spl72_9 | spl72_10 | spl72_11 | spl72_12), inference(avatar_split_clause, [], [f181, f777, f773, f769, f765])).
fof(f181, plain, ((e13 = op1(e13, e11)) | (e12 = op1(e13, e11)) | (e11 = op1(e13, e11)) | (e10 = op1(e13, e11))), inference(cnf_transformation, [], [f1])).
fof(f763, plain, (spl72_5 | spl72_6 | spl72_7 | spl72_8), inference(avatar_split_clause, [], [f182, f760, f756, f752, f748])).
fof(f182, plain, ((e13 = op1(e13, e12)) | (e12 = op1(e13, e12)) | (e11 = op1(e13, e12)) | (e10 = op1(e13, e12))), inference(cnf_transformation, [], [f1])).
fof(f746, plain, (spl72_1 | spl72_2 | spl72_3 | spl72_4), inference(avatar_split_clause, [], [f183, f743, f739, f735, f731])).
fof(f183, plain, ((e13 = op1(e13, e13)) | (e12 = op1(e13, e13)) | (e11 = op1(e13, e13)) | (e10 = op1(e13, e13))), inference(cnf_transformation, [], [f1])).