fof(f6599, plain, $false, inference(avatar_sat_refutation, [], [f1040, f1057, f1074, f1108, f1125, f1142, f1159, f1176, f1193, f1210, f1227, f1261, f1297, f1298, f1300, f1301, f1302, f1303, f1304, f1305, f1306, f1309, f1310, f1311, f1315, f1316, f1317, f1318, f1319, f1320, f1326, f1327, f1361, f1378, f1429, f1446, f1463, f1480, f1497, f1514, f1531, f1548, f1565, f1601, f1602, f1604, f1606, f1607, f1608, f1610, f1611, f1612, f1617, f1619, f1620, f1622, f1623, f1624, f1630, f1631, f1640, f1648, f1649, f1655, f1658, f1665, f1667, f1673, f1676, f1682, f1685, f1694, f1703, f1710, f1712, f1719, f1721, f1726, f1730, f1739, f1747, f1748, f1753, f1766, f1780, f1789, f1794, f1803, f1808, f1818, f1828, f1837, f1842, f1856, f1865, f1866, f1875, f1876, f1886, f1895, f1900, f1910, f1920, f1930, f1940, f1947, f1948, f1949, f1960, f1970, f1979, f1986, f1990, f1995, f1997, f1998, f1999, f2000, f2010, f2016, f2017, f2019, f2020, f2027, f2029, f2030, f2037, f2040, f2050, f2059, f2069, f2070, f2080, f2090, f2091, f2095, f2101, f2102, f2104, f2107, f2116, f2124, f2125, f2131, f2134, f2141, f2143, f2149, f2152, f2158, f2161, f2170, f2179, f2186, f2188, f2195, f2197, f2202, f2206, f2212, f2215, f2223, f2224, f2229, f2242, f2256, f2265, f2279, f2294, f2304, f2313, f2318, f2332, f2341, f2342, f2351, f2352, f2362, f2371, f2376, f2386, f2396, f2406, f2416, f2421, f2424, f2425, f2436, f2446, f2455, f2462, f2464, f2466, f2471, f2473, f2474, f2475, f2486, f2492, f2493, f2495, f2496, f2502, f2505, f2506, f2512, f2513, f2516, f2526, f2535, f2545, f2546, f2556, f2566, f2567, f2571, f2577, f2578, f2581, f2583, f2585, f2586, f2588, f2589, f2641, f2666, f2676, f2979, f3338, f3356, f3358, f3381, f3382, f3388, f3393, f3394, f3415, f3430, f3449, f3456, f3470, f3474, f3482, f3483, f3484, f3485, f3486, f3487, f3488, f3503, f3505, f3528, f3529, f3536, f3538, f3557, f3558, f3564, f3565, f3568, f3578, f3586, f3594, f3597, f3606, f3612, f3622, f3646, f3648, f3656, f3671, f3693, f3694, f3695, f3696, f3703, f3722, f3729, f3731, f3743, f3750, f3753, f3775, f3777, f3780, f3785, f3804, f3808, f3832, f3838, f3839, f3855, f3890, f3901, f3919, f3929, f3931, f3942, f3944, f3949, f3951, f3964, f3966, f3970, f3993, f4016, f4020, f4038, f4045, f4054, f4059, f4081, f4093, f4104, f4105, f4119, f4123, f4134, f4156, f4172, f4186, f4206, f4214, f4220, f4223, f4239, f4240, f4244, f4250, f4258, f4266, f4268, f4298, f4312, f4314, f4323, f4354, f4366, f4376, f4379, f4385, f4394, f4404, f4406, f4449, f4469, f4518, f4528, f4548, f4560, f4590, f4617, f4620, f4664, f4687, f4702, f4705, f4728, f4752, f4776, f4779, f4786, f4787, f4794, f4804, f4823, f4826, f4833, f4861, f4862, f4877, f4897, f4928, f4966, f4990, f5038, f5072, f5079, f5086, f5088, f5091, f5150, f5152, f5155, f5197, f5204, f5227, f5239, f5244, f5257, f5282, f5306, f5333, f5359, f5383, f5395, f5404, f5419, f5439, f5442, f5443, f5466, f5494, f5513, f5527, f5538, f5542, f5567, f5570, f5575, f5588, f5605, f5624, f5628, f5675, f5676, f5678, f5711, f5753, f5814, f5816, f5819, f5821, f5826, f5900, f5923, f5924, f5969, f5970, f5972, f5975, f5986, f5990, f5992, f6000, f6002, f6012, f6061, f6062, f6076, f6090, f6091, f6101, f6111, f6216, f6293, f6300, f6360, f6413, f6423, f6446, f6464, f6465, f6477, f6488, f6507, f6508, f6510, f6511, f6526, f6590, f6598])).
fof(f6598, plain, (~ spl102_5 | spl102_17 | ~ spl102_146), inference(avatar_contradiction_clause, [], [f6597])).
fof(f6597, plain, ($false | (~ spl102_5 | spl102_17 | ~ spl102_146)), inference(subsumption_resolution, [], [f6596, f1094])).
fof(f1094, plain, (~ (e10 = op1(e12, e13)) | spl102_17), inference(avatar_component_clause, [], [f1093])).
fof(f1093, plain, (spl102_17 <=> (e10 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl102_17])])).
fof(f6596, plain, ((e10 = op1(e12, e13)) | (~ spl102_5 | ~ spl102_146)), inference(forward_demodulation, [], [f1779, f1044])).
fof(f1044, plain, ((e10 = op1(e13, e12)) | ~ spl102_5), inference(avatar_component_clause, [], [f1042])).
fof(f1042, plain, (spl102_5 <=> (e10 = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl102_5])])).
fof(f1779, plain, ((op1(e12, e13) = op1(e13, e12)) | ~ spl102_146), inference(avatar_component_clause, [], [f1777])).
fof(f1777, plain, (spl102_146 <=> (op1(e12, e13) = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl102_146])])).
fof(f6590, plain, (~ spl102_10 | spl102_42 | ~ spl102_148), inference(avatar_contradiction_clause, [], [f6589])).
fof(f6589, plain, ($false | (~ spl102_10 | spl102_42 | ~ spl102_148)), inference(subsumption_resolution, [], [f6588, f1200])).
fof(f1200, plain, (~ (e11 = op1(e11, e11)) | spl102_42), inference(avatar_component_clause, [], [f1199])).
fof(f1199, plain, (spl102_42 <=> (e11 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl102_42])])).
fof(f6588, plain, ((e11 = op1(e11, e11)) | (~ spl102_10 | ~ spl102_148)), inference(forward_demodulation, [], [f1788, f1065])).
fof(f1065, plain, ((e11 = op1(e13, e11)) | ~ spl102_10), inference(avatar_component_clause, [], [f1063])).
fof(f1063, plain, (spl102_10 <=> (e11 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl102_10])])).
fof(f1788, plain, ((e11 = op1(e11, op1(e13, e11))) | ~ spl102_148), inference(avatar_component_clause, [], [f1786])).
fof(f1786, plain, (spl102_148 <=> (e11 = op1(e11, op1(e13, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl102_148])])).
fof(f6526, plain, (~ spl102_17 | ~ spl102_25), inference(avatar_split_clause, [], [f6518, f1127, f1093])).
fof(f1127, plain, (spl102_25 <=> (e10 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl102_25])])).
fof(f6518, plain, (~ (e10 = op1(e12, e13)) | ~ spl102_25), inference(backward_demodulation, [], [f366, f1129])).
fof(f1129, plain, ((e10 = op1(e12, e11)) | ~ spl102_25), inference(avatar_component_clause, [], [f1127])).
fof(f366, plain, ~ (op1(e12, e11) = op1(e12, e13)), inference(cnf_transformation, [], [f5])).
fof(f5, plain, (~ (op1(e13, e12) = op1(e13, e13)) & ~ (op1(e13, e11) = op1(e13, e13)) & ~ (op1(e13, e11) = op1(e13, e12)) & ~ (op1(e13, e10) = op1(e13, e13)) & ~ (op1(e13, e10) = op1(e13, e12)) & ~ (op1(e13, e10) = op1(e13, e11)) & ~ (op1(e12, e12) = op1(e12, e13)) & ~ (op1(e12, e11) = op1(e12, e13)) & ~ (op1(e12, e11) = op1(e12, e12)) & ~ (op1(e12, e10) = op1(e12, e13)) & ~ (op1(e12, e10) = op1(e12, e12)) & ~ (op1(e12, e10) = op1(e12, e11)) & ~ (op1(e11, e12) = op1(e11, e13)) & ~ (op1(e11, e11) = op1(e11, e13)) & ~ (op1(e11, e11) = op1(e11, e12)) & ~ (op1(e11, e10) = op1(e11, e13)) & ~ (op1(e11, e10) = op1(e11, e12)) & ~ (op1(e11, e10) = op1(e11, e11)) & ~ (op1(e10, e12) = op1(e10, e13)) & ~ (op1(e10, e11) = op1(e10, e13)) & ~ (op1(e10, e11) = op1(e10, e12)) & ~ (op1(e10, e10) = op1(e10, e13)) & ~ (op1(e10, e10) = op1(e10, e12)) & ~ (op1(e10, e10) = op1(e10, e11)) & ~ (op1(e12, e13) = op1(e13, e13)) & ~ (op1(e11, e13) = op1(e13, e13)) & ~ (op1(e11, e13) = op1(e12, e13)) & ~ (op1(e10, e13) = op1(e13, e13)) & ~ (op1(e10, e13) = op1(e12, e13)) & ~ (op1(e10, e13) = op1(e11, e13)) & ~ (op1(e12, e12) = op1(e13, e12)) & ~ (op1(e11, e12) = op1(e13, e12)) & ~ (op1(e11, e12) = op1(e12, e12)) & ~ (op1(e10, e12) = op1(e13, e12)) & ~ (op1(e10, e12) = op1(e12, e12)) & ~ (op1(e10, e12) = op1(e11, e12)) & ~ (op1(e12, e11) = op1(e13, e11)) & ~ (op1(e11, e11) = op1(e13, e11)) & ~ (op1(e11, e11) = op1(e12, e11)) & ~ (op1(e10, e11) = op1(e13, e11)) & ~ (op1(e10, e11) = op1(e12, e11)) & ~ (op1(e10, e11) = op1(e11, e11)) & ~ (op1(e12, e10) = op1(e13, e10)) & ~ (op1(e11, e10) = op1(e13, e10)) & ~ (op1(e11, e10) = op1(e12, e10)) & ~ (op1(e10, e10) = op1(e13, e10)) & ~ (op1(e10, e10) = op1(e12, e10)) & ~ (op1(e10, e10) = op1(e11, e10))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG109+1.p', ax5)).
fof(f6511, plain, (~ spl102_26 | ~ spl102_34 | ~ spl102_149), inference(avatar_split_clause, [], [f6502, f1791, f1165, f1131])).
fof(f1131, plain, (spl102_26 <=> (e11 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl102_26])])).
fof(f1165, plain, (spl102_34 <=> (e11 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl102_34])])).
fof(f1791, plain, (spl102_149 <=> (op1(e11, e13) = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl102_149])])).
fof(f6502, plain, (~ (e11 = op1(e12, e11)) | (~ spl102_34 | ~ spl102_149)), inference(backward_demodulation, [], [f6431, f1167])).
fof(f1167, plain, ((e11 = op1(e11, e13)) | ~ spl102_34), inference(avatar_component_clause, [], [f1165])).
fof(f6431, plain, (~ (op1(e11, e13) = op1(e12, e11)) | ~ spl102_149), inference(backward_demodulation, [], [f337, f1793])).
fof(f1793, plain, ((op1(e11, e13) = op1(e13, e11)) | ~ spl102_149), inference(avatar_component_clause, [], [f1791])).
fof(f337, plain, ~ (op1(e12, e11) = op1(e13, e11)), inference(cnf_transformation, [], [f5])).
fof(f6510, plain, (spl102_12 | ~ spl102_34 | ~ spl102_164), inference(avatar_contradiction_clause, [], [f6509])).
fof(f6509, plain, ($false | (spl102_12 | ~ spl102_34 | ~ spl102_164)), inference(subsumption_resolution, [], [f6498, f1072])).
fof(f1072, plain, (~ (e13 = op1(e13, e11)) | spl102_12), inference(avatar_component_clause, [], [f1071])).
fof(f1071, plain, (spl102_12 <=> (e13 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl102_12])])).
fof(f6498, plain, ((e13 = op1(e13, e11)) | (~ spl102_34 | ~ spl102_164)), inference(backward_demodulation, [], [f1864, f1167])).
fof(f1864, plain, ((e13 = op1(e13, op1(e11, e13))) | ~ spl102_164), inference(avatar_component_clause, [], [f1862])).
fof(f1862, plain, (spl102_164 <=> (e13 = op1(e13, op1(e11, e13)))), introduced(avatar_definition, [new_symbols(naming, [spl102_164])])).
fof(f6508, plain, (spl102_10 | ~ spl102_34 | ~ spl102_149), inference(avatar_split_clause, [], [f6497, f1791, f1165, f1063])).
fof(f6497, plain, ((e11 = op1(e13, e11)) | (~ spl102_34 | ~ spl102_149)), inference(backward_demodulation, [], [f1793, f1167])).
fof(f6507, plain, (~ spl102_2 | ~ spl102_34), inference(avatar_split_clause, [], [f6494, f1165, f1029])).
fof(f1029, plain, (spl102_2 <=> (e11 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl102_2])])).
fof(f6494, plain, (~ (e11 = op1(e13, e13)) | ~ spl102_34), inference(backward_demodulation, [], [f348, f1167])).
fof(f348, plain, ~ (op1(e11, e13) = op1(e13, e13)), inference(cnf_transformation, [], [f5])).
fof(f6488, plain, (~ spl102_8 | ~ spl102_40), inference(avatar_split_clause, [], [f6480, f1190, f1054])).
fof(f1054, plain, (spl102_8 <=> (e13 = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl102_8])])).
fof(f1190, plain, (spl102_40 <=> (e13 = op1(e11, e12))), introduced(avatar_definition, [new_symbols(naming, [spl102_40])])).
fof(f6480, plain, (~ (e13 = op1(e13, e12)) | ~ spl102_40), inference(backward_demodulation, [], [f342, f1192])).
fof(f1192, plain, ((e13 = op1(e11, e12)) | ~ spl102_40), inference(avatar_component_clause, [], [f1190])).
fof(f342, plain, ~ (op1(e11, e12) = op1(e13, e12)), inference(cnf_transformation, [], [f5])).
fof(f6477, plain, (~ spl102_17 | ~ spl102_49), inference(avatar_split_clause, [], [f6469, f1229, f1093])).
fof(f1229, plain, (spl102_49 <=> (e10 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl102_49])])).
fof(f6469, plain, (~ (e10 = op1(e12, e13)) | ~ spl102_49), inference(backward_demodulation, [], [f345, f1231])).
fof(f1231, plain, ((e10 = op1(e10, e13)) | ~ spl102_49), inference(avatar_component_clause, [], [f1229])).
fof(f345, plain, ~ (op1(e10, e13) = op1(e12, e13)), inference(cnf_transformation, [], [f5])).
fof(f6465, plain, (~ spl102_51 | ~ spl102_55), inference(avatar_split_clause, [], [f6457, f1254, f1237])).
fof(f1237, plain, (spl102_51 <=> (e12 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl102_51])])).
fof(f1254, plain, (spl102_55 <=> (e12 = op1(e10, e12))), introduced(avatar_definition, [new_symbols(naming, [spl102_55])])).
fof(f6457, plain, (~ (e12 = op1(e10, e13)) | ~ spl102_55), inference(backward_demodulation, [], [f355, f1256])).
fof(f1256, plain, ((e12 = op1(e10, e12)) | ~ spl102_55), inference(avatar_component_clause, [], [f1254])).
fof(f355, plain, ~ (op1(e10, e12) = op1(e10, e13)), inference(cnf_transformation, [], [f5])).
fof(f6464, plain, (~ spl102_23 | ~ spl102_55), inference(avatar_split_clause, [], [f6455, f1254, f1118])).
fof(f1118, plain, (spl102_23 <=> (e12 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl102_23])])).
fof(f6455, plain, (~ (e12 = op1(e12, e12)) | ~ spl102_55), inference(backward_demodulation, [], [f339, f1256])).
fof(f339, plain, ~ (op1(e10, e12) = op1(e12, e12)), inference(cnf_transformation, [], [f5])).
fof(f6446, plain, (~ spl102_113 | ~ spl102_117), inference(avatar_split_clause, [], [f6439, f1550, f1533])).
fof(f1533, plain, (spl102_113 <=> (e20 = op2(e20, e23))), introduced(avatar_definition, [new_symbols(naming, [spl102_113])])).
fof(f1550, plain, (spl102_117 <=> (e20 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl102_117])])).
fof(f6439, plain, (~ (e20 = op2(e20, e23)) | ~ spl102_117), inference(backward_demodulation, [], [f403, f1552])).
fof(f1552, plain, ((e20 = op2(e20, e22)) | ~ spl102_117), inference(avatar_component_clause, [], [f1550])).
fof(f403, plain, ~ (op2(e20, e22) = op2(e20, e23)), inference(cnf_transformation, [], [f6])).
fof(f6, plain, (~ (op2(e23, e22) = op2(e23, e23)) & ~ (op2(e23, e21) = op2(e23, e23)) & ~ (op2(e23, e21) = op2(e23, e22)) & ~ (op2(e23, e20) = op2(e23, e23)) & ~ (op2(e23, e20) = op2(e23, e22)) & ~ (op2(e23, e20) = op2(e23, e21)) & ~ (op2(e22, e22) = op2(e22, e23)) & ~ (op2(e22, e21) = op2(e22, e23)) & ~ (op2(e22, e21) = op2(e22, e22)) & ~ (op2(e22, e20) = op2(e22, e23)) & ~ (op2(e22, e20) = op2(e22, e22)) & ~ (op2(e22, e20) = op2(e22, e21)) & ~ (op2(e21, e22) = op2(e21, e23)) & ~ (op2(e21, e21) = op2(e21, e23)) & ~ (op2(e21, e21) = op2(e21, e22)) & ~ (op2(e21, e20) = op2(e21, e23)) & ~ (op2(e21, e20) = op2(e21, e22)) & ~ (op2(e21, e20) = op2(e21, e21)) & ~ (op2(e20, e22) = op2(e20, e23)) & ~ (op2(e20, e21) = op2(e20, e23)) & ~ (op2(e20, e21) = op2(e20, e22)) & ~ (op2(e20, e20) = op2(e20, e23)) & ~ (op2(e20, e20) = op2(e20, e22)) & ~ (op2(e20, e20) = op2(e20, e21)) & ~ (op2(e22, e23) = op2(e23, e23)) & ~ (op2(e21, e23) = op2(e23, e23)) & ~ (op2(e21, e23) = op2(e22, e23)) & ~ (op2(e20, e23) = op2(e23, e23)) & ~ (op2(e20, e23) = op2(e22, e23)) & ~ (op2(e20, e23) = op2(e21, e23)) & ~ (op2(e22, e22) = op2(e23, e22)) & ~ (op2(e21, e22) = op2(e23, e22)) & ~ (op2(e21, e22) = op2(e22, e22)) & ~ (op2(e20, e22) = op2(e23, e22)) & ~ (op2(e20, e22) = op2(e22, e22)) & ~ (op2(e20, e22) = op2(e21, e22)) & ~ (op2(e22, e21) = op2(e23, e21)) & ~ (op2(e21, e21) = op2(e23, e21)) & ~ (op2(e21, e21) = op2(e22, e21)) & ~ (op2(e20, e21) = op2(e23, e21)) & ~ (op2(e20, e21) = op2(e22, e21)) & ~ (op2(e20, e21) = op2(e21, e21)) & ~ (op2(e22, e20) = op2(e23, e20)) & ~ (op2(e21, e20) = op2(e23, e20)) & ~ (op2(e21, e20) = op2(e22, e20)) & ~ (op2(e20, e20) = op2(e23, e20)) & ~ (op2(e20, e20) = op2(e22, e20)) & ~ (op2(e20, e20) = op2(e21, e20))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG109+1.p', ax6)).
fof(f6423, plain, (spl102_117 | ~ spl102_79 | ~ spl102_218), inference(avatar_split_clause, [], [f6422, f2276, f1388, f1550])).
fof(f1388, plain, (spl102_79 <=> (e22 = op2(e23, e20))), introduced(avatar_definition, [new_symbols(naming, [spl102_79])])).
fof(f2276, plain, (spl102_218 <=> (e20 = op2(e20, op2(e23, e20)))), introduced(avatar_definition, [new_symbols(naming, [spl102_218])])).
fof(f6422, plain, ((e20 = op2(e20, e22)) | (~ spl102_79 | ~ spl102_218)), inference(forward_demodulation, [], [f2278, f1390])).
fof(f1390, plain, ((e22 = op2(e23, e20)) | ~ spl102_79), inference(avatar_component_clause, [], [f1388])).
fof(f2278, plain, ((e20 = op2(e20, op2(e23, e20))) | ~ spl102_218), inference(avatar_component_clause, [], [f2276])).
fof(f6413, plain, (~ spl102_119 | ~ spl102_278), inference(avatar_split_clause, [], [f6401, f2703, f1558])).
fof(f1558, plain, (spl102_119 <=> (e22 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl102_119])])).
fof(f2703, plain, (spl102_278 <=> (e22 = h3(e11))), introduced(avatar_definition, [new_symbols(naming, [spl102_278])])).
fof(f6401, plain, (~ (e22 = op2(e20, e22)) | ~ spl102_278), inference(backward_demodulation, [], [f2614, f2704])).
fof(f2704, plain, ((e22 = h3(e11)) | ~ spl102_278), inference(avatar_component_clause, [], [f2703])).
fof(f2614, plain, ~ (op2(e20, e22) = h3(e11)), inference(backward_demodulation, [], [f387, f913])).
fof(f913, plain, (op2(e22, e22) = h3(e11)), inference(cnf_transformation, [], [f16])).
fof(f16, plain, ((op2(e22, op2(e22, e22)) = h3(e13)) & (h3(e12) = op2(op2(e22, op2(e22, e22)), e22)) & (op2(e22, e22) = h3(e11)) & (e22 = h3(e10))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG109+1.p', ax16)).
fof(f387, plain, ~ (op2(e20, e22) = op2(e22, e22)), inference(cnf_transformation, [], [f6])).
fof(f6360, plain, (~ spl102_73 | ~ spl102_109 | ~ spl102_215), inference(avatar_contradiction_clause, [], [f6359])).
fof(f6359, plain, ($false | (~ spl102_73 | ~ spl102_109 | ~ spl102_215)), inference(subsumption_resolution, [], [f6358, f428])).
fof(f428, plain, ~ (e20 = e21), inference(cnf_transformation, [], [f8])).
fof(f8, plain, (~ (e22 = e23) & ~ (e21 = e23) & ~ (e21 = e22) & ~ (e20 = e23) & ~ (e20 = e22) & ~ (e20 = e21)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG109+1.p', ax8)).
fof(f6358, plain, ((e20 = e21) | (~ spl102_73 | ~ spl102_109 | ~ spl102_215)), inference(forward_demodulation, [], [f6357, f1518])).
fof(f1518, plain, ((e20 = op2(e21, e20)) | ~ spl102_109), inference(avatar_component_clause, [], [f1516])).
fof(f1516, plain, (spl102_109 <=> (e20 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl102_109])])).
fof(f6357, plain, ((e21 = op2(e21, e20)) | (~ spl102_73 | ~ spl102_215)), inference(forward_demodulation, [], [f2264, f1365])).
fof(f1365, plain, ((e20 = op2(e23, e21)) | ~ spl102_73), inference(avatar_component_clause, [], [f1363])).
fof(f1363, plain, (spl102_73 <=> (e20 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl102_73])])).
fof(f2264, plain, ((e21 = op2(e21, op2(e23, e21))) | ~ spl102_215), inference(avatar_component_clause, [], [f2262])).
fof(f2262, plain, (spl102_215 <=> (e21 = op2(e21, op2(e23, e21)))), introduced(avatar_definition, [new_symbols(naming, [spl102_215])])).
fof(f6300, plain, (spl102_76 | ~ spl102_100 | ~ spl102_216), inference(avatar_contradiction_clause, [], [f6299])).
fof(f6299, plain, ($false | (spl102_76 | ~ spl102_100 | ~ spl102_216)), inference(subsumption_resolution, [], [f6295, f1376])).
fof(f1376, plain, (~ (e23 = op2(e23, e21)) | spl102_76), inference(avatar_component_clause, [], [f1375])).
fof(f1375, plain, (spl102_76 <=> (e23 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl102_76])])).
fof(f6295, plain, ((e23 = op2(e23, e21)) | (~ spl102_100 | ~ spl102_216)), inference(backward_demodulation, [], [f2269, f1479])).
fof(f1479, plain, ((e23 = op2(e21, e23)) | ~ spl102_100), inference(avatar_component_clause, [], [f1477])).
fof(f1477, plain, (spl102_100 <=> (e23 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl102_100])])).
fof(f2269, plain, ((op2(e21, e23) = op2(e23, e21)) | ~ spl102_216), inference(avatar_component_clause, [], [f2267])).
fof(f2267, plain, (spl102_216 <=> (op2(e21, e23) = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl102_216])])).
fof(f6293, plain, (~ spl102_93 | ~ spl102_109), inference(avatar_split_clause, [], [f6289, f1516, f1448])).
fof(f1448, plain, (spl102_93 <=> (e20 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl102_93])])).
fof(f6289, plain, (~ (e20 = op2(e22, e20)) | ~ spl102_109), inference(backward_demodulation, [], [f377, f1518])).
fof(f377, plain, ~ (op2(e21, e20) = op2(e22, e20)), inference(cnf_transformation, [], [f6])).
fof(f6216, plain, (~ spl102_99 | ~ spl102_290), inference(avatar_split_clause, [], [f6210, f2764, f1473])).
fof(f1473, plain, (spl102_99 <=> (e22 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl102_99])])).
fof(f2764, plain, (spl102_290 <=> (e22 = h2(e11))), introduced(avatar_definition, [new_symbols(naming, [spl102_290])])).
fof(f6210, plain, (~ (e22 = op2(e21, e23)) | ~ spl102_290), inference(backward_demodulation, [], [f2611, f2765])).
fof(f2765, plain, ((e22 = h2(e11)) | ~ spl102_290), inference(avatar_component_clause, [], [f2764])).
fof(f2611, plain, ~ (op2(e21, e23) = h2(e11)), inference(backward_demodulation, [], [f408, f909])).
fof(f909, plain, (op2(e21, e21) = h2(e11)), inference(cnf_transformation, [], [f15])).
fof(f15, plain, ((op2(e21, op2(e21, e21)) = h2(e13)) & (h2(e12) = op2(op2(e21, op2(e21, e21)), e21)) & (op2(e21, e21) = h2(e11)) & (e21 = h2(e10))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG109+1.p', ax15)).
fof(f408, plain, ~ (op2(e21, e21) = op2(e21, e23)), inference(cnf_transformation, [], [f6])).
fof(f6111, plain, (~ spl102_19 | ~ spl102_23), inference(avatar_split_clause, [], [f6107, f1118, f1101])).
fof(f1101, plain, (spl102_19 <=> (e12 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl102_19])])).
fof(f6107, plain, (~ (e12 = op1(e12, e13)) | ~ spl102_23), inference(backward_demodulation, [], [f367, f1120])).
fof(f1120, plain, ((e12 = op1(e12, e12)) | ~ spl102_23), inference(avatar_component_clause, [], [f1118])).
fof(f367, plain, ~ (op1(e12, e12) = op1(e12, e13)), inference(cnf_transformation, [], [f5])).
fof(f6101, plain, (~ spl102_4 | ~ spl102_36), inference(avatar_split_clause, [], [f6098, f1173, f1037])).
fof(f1037, plain, (spl102_4 <=> (e13 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl102_4])])).
fof(f1173, plain, (spl102_36 <=> (e13 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl102_36])])).
fof(f6098, plain, (~ (e13 = op1(e13, e13)) | ~ spl102_36), inference(backward_demodulation, [], [f348, f1175])).
fof(f1175, plain, ((e13 = op1(e11, e13)) | ~ spl102_36), inference(avatar_component_clause, [], [f1173])).
fof(f6091, plain, (~ spl102_34 | ~ spl102_38), inference(avatar_split_clause, [], [f6083, f1182, f1165])).
fof(f1182, plain, (spl102_38 <=> (e11 = op1(e11, e12))), introduced(avatar_definition, [new_symbols(naming, [spl102_38])])).
fof(f6083, plain, (~ (e11 = op1(e11, e13)) | ~ spl102_38), inference(backward_demodulation, [], [f361, f1184])).
fof(f1184, plain, ((e11 = op1(e11, e12)) | ~ spl102_38), inference(avatar_component_clause, [], [f1182])).
fof(f361, plain, ~ (op1(e11, e12) = op1(e11, e13)), inference(cnf_transformation, [], [f5])).
fof(f6090, plain, (~ spl102_6 | ~ spl102_38), inference(avatar_split_clause, [], [f6082, f1182, f1046])).
fof(f1046, plain, (spl102_6 <=> (e11 = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl102_6])])).
fof(f6082, plain, (~ (e11 = op1(e13, e12)) | ~ spl102_38), inference(backward_demodulation, [], [f342, f1184])).
fof(f6076, plain, (~ spl102_19 | ~ spl102_51), inference(avatar_split_clause, [], [f6070, f1237, f1101])).
fof(f6070, plain, (~ (e12 = op1(e12, e13)) | ~ spl102_51), inference(backward_demodulation, [], [f345, f1239])).
fof(f1239, plain, ((e12 = op1(e10, e13)) | ~ spl102_51), inference(avatar_component_clause, [], [f1237])).
fof(f6062, plain, (~ spl102_49 | ~ spl102_53), inference(avatar_split_clause, [], [f6053, f1246, f1229])).
fof(f1246, plain, (spl102_53 <=> (e10 = op1(e10, e12))), introduced(avatar_definition, [new_symbols(naming, [spl102_53])])).
fof(f6053, plain, (~ (e10 = op1(e10, e13)) | ~ spl102_53), inference(backward_demodulation, [], [f355, f1248])).
fof(f1248, plain, ((e10 = op1(e10, e12)) | ~ spl102_53), inference(avatar_component_clause, [], [f1246])).
fof(f6061, plain, (~ spl102_21 | ~ spl102_53), inference(avatar_split_clause, [], [f6051, f1246, f1110])).
fof(f1110, plain, (spl102_21 <=> (e10 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl102_21])])).
fof(f6051, plain, (~ (e10 = op1(e12, e12)) | ~ spl102_53), inference(backward_demodulation, [], [f339, f1248])).
fof(f6012, plain, (~ spl102_100 | ~ spl102_112), inference(avatar_split_clause, [], [f6011, f1528, f1477])).
fof(f1528, plain, (spl102_112 <=> (e23 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl102_112])])).
fof(f6011, plain, (~ (e23 = op2(e21, e23)) | ~ spl102_112), inference(backward_demodulation, [], [f406, f1530])).
fof(f1530, plain, ((e23 = op2(e21, e20)) | ~ spl102_112), inference(avatar_component_clause, [], [f1528])).
fof(f406, plain, ~ (op2(e21, e20) = op2(e21, e23)), inference(cnf_transformation, [], [f6])).
fof(f6002, plain, (spl102_95 | ~ spl102_119 | ~ spl102_229), inference(avatar_contradiction_clause, [], [f6001])).
fof(f6001, plain, ($false | (spl102_95 | ~ spl102_119 | ~ spl102_229)), inference(subsumption_resolution, [], [f5996, f1457])).
fof(f1457, plain, (~ (e22 = op2(e22, e20)) | spl102_95), inference(avatar_component_clause, [], [f1456])).
fof(f1456, plain, (spl102_95 <=> (e22 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl102_95])])).
fof(f5996, plain, ((e22 = op2(e22, e20)) | (~ spl102_119 | ~ spl102_229)), inference(backward_demodulation, [], [f2331, f1560])).
fof(f1560, plain, ((e22 = op2(e20, e22)) | ~ spl102_119), inference(avatar_component_clause, [], [f1558])).
fof(f2331, plain, ((op2(e20, e22) = op2(e22, e20)) | ~ spl102_229), inference(avatar_component_clause, [], [f2329])).
fof(f2329, plain, (spl102_229 <=> (op2(e20, e22) = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl102_229])])).
fof(f6000, plain, (~ spl102_115 | ~ spl102_119), inference(avatar_split_clause, [], [f5994, f1558, f1541])).
fof(f1541, plain, (spl102_115 <=> (e22 = op2(e20, e23))), introduced(avatar_definition, [new_symbols(naming, [spl102_115])])).
fof(f5994, plain, (~ (e22 = op2(e20, e23)) | ~ spl102_119), inference(backward_demodulation, [], [f403, f1560])).
fof(f5992, plain, (spl102_53 | ~ spl102_15 | ~ spl102_151), inference(avatar_split_clause, [], [f5991, f1800, f1084, f1246])).
fof(f1084, plain, (spl102_15 <=> (e12 = op1(e13, e10))), introduced(avatar_definition, [new_symbols(naming, [spl102_15])])).
fof(f1800, plain, (spl102_151 <=> (e10 = op1(e10, op1(e13, e10)))), introduced(avatar_definition, [new_symbols(naming, [spl102_151])])).
fof(f5991, plain, ((e10 = op1(e10, e12)) | (~ spl102_15 | ~ spl102_151)), inference(forward_demodulation, [], [f1802, f1086])).
fof(f1086, plain, ((e12 = op1(e13, e10)) | ~ spl102_15), inference(avatar_component_clause, [], [f1084])).
fof(f1802, plain, ((e10 = op1(e10, op1(e13, e10))) | ~ spl102_151), inference(avatar_component_clause, [], [f1800])).
fof(f5990, plain, (spl102_51 | ~ spl102_15 | ~ spl102_152), inference(avatar_split_clause, [], [f5989, f1805, f1084, f1237])).
fof(f1805, plain, (spl102_152 <=> (op1(e10, e13) = op1(e13, e10))), introduced(avatar_definition, [new_symbols(naming, [spl102_152])])).
fof(f5989, plain, ((e12 = op1(e10, e13)) | (~ spl102_15 | ~ spl102_152)), inference(forward_demodulation, [], [f1807, f1086])).
fof(f1807, plain, ((op1(e10, e13) = op1(e13, e10)) | ~ spl102_152), inference(avatar_component_clause, [], [f1805])).
fof(f5986, plain, (~ spl102_32 | spl102_56 | ~ spl102_162), inference(avatar_contradiction_clause, [], [f5985])).
fof(f5985, plain, ($false | (~ spl102_32 | spl102_56 | ~ spl102_162)), inference(subsumption_resolution, [], [f5984, f1259])).
fof(f1259, plain, (~ (e13 = op1(e10, e12)) | spl102_56), inference(avatar_component_clause, [], [f1258])).
fof(f1258, plain, (spl102_56 <=> (e13 = op1(e10, e12))), introduced(avatar_definition, [new_symbols(naming, [spl102_56])])).
fof(f5984, plain, ((e13 = op1(e10, e12)) | (~ spl102_32 | ~ spl102_162)), inference(forward_demodulation, [], [f1855, f1158])).
fof(f1158, plain, ((e13 = op1(e12, e10)) | ~ spl102_32), inference(avatar_component_clause, [], [f1156])).
fof(f1156, plain, (spl102_32 <=> (e13 = op1(e12, e10))), introduced(avatar_definition, [new_symbols(naming, [spl102_32])])).
fof(f1855, plain, ((op1(e10, e12) = op1(e12, e10)) | ~ spl102_162), inference(avatar_component_clause, [], [f1853])).
fof(f1853, plain, (spl102_162 <=> (op1(e10, e12) = op1(e12, e10))), introduced(avatar_definition, [new_symbols(naming, [spl102_162])])).
fof(f5975, plain, (spl102_112 | ~ spl102_124 | ~ spl102_238), inference(avatar_split_clause, [], [f5974, f2373, f1579, f1528])).
fof(f1579, plain, (spl102_124 <=> (e23 = op2(e20, e21))), introduced(avatar_definition, [new_symbols(naming, [spl102_124])])).
fof(f2373, plain, (spl102_238 <=> (op2(e20, e21) = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl102_238])])).
fof(f5974, plain, ((e23 = op2(e21, e20)) | (~ spl102_124 | ~ spl102_238)), inference(forward_demodulation, [], [f2375, f1581])).
fof(f1581, plain, ((e23 = op2(e20, e21)) | ~ spl102_124), inference(avatar_component_clause, [], [f1579])).
fof(f2375, plain, ((op2(e20, e21) = op2(e21, e20)) | ~ spl102_238), inference(avatar_component_clause, [], [f2373])).
fof(f5972, plain, (spl102_110 | ~ spl102_292 | ~ spl102_298), inference(avatar_contradiction_clause, [], [f5971])).
fof(f5971, plain, ($false | (spl102_110 | ~ spl102_292 | ~ spl102_298)), inference(subsumption_resolution, [], [f5967, f1521])).
fof(f1521, plain, (~ (e21 = op2(e21, e20)) | spl102_110), inference(avatar_component_clause, [], [f1520])).
fof(f1520, plain, (spl102_110 <=> (e21 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl102_110])])).
fof(f5967, plain, ((e21 = op2(e21, e20)) | (~ spl102_292 | ~ spl102_298)), inference(backward_demodulation, [], [f5874, f2806])).
fof(f2806, plain, ((e20 = h2(e11)) | ~ spl102_298), inference(avatar_component_clause, [], [f2805])).
fof(f2805, plain, (spl102_298 <=> (e20 = h2(e11))), introduced(avatar_definition, [new_symbols(naming, [spl102_298])])).
fof(f5874, plain, ((e21 = op2(e21, h2(e11))) | ~ spl102_292), inference(forward_demodulation, [], [f2612, f2775])).
fof(f2775, plain, ((e21 = h2(e13)) | ~ spl102_292), inference(avatar_component_clause, [], [f2774])).
fof(f2774, plain, (spl102_292 <=> (e21 = h2(e13))), introduced(avatar_definition, [new_symbols(naming, [spl102_292])])).
fof(f2612, plain, (h2(e13) = op2(e21, h2(e11))), inference(backward_demodulation, [], [f911, f909])).
fof(f911, plain, (op2(e21, op2(e21, e21)) = h2(e13)), inference(cnf_transformation, [], [f15])).
fof(f5970, plain, (~ spl102_109 | ~ spl102_298), inference(avatar_split_clause, [], [f5963, f2805, f1516])).
fof(f5963, plain, (~ (e20 = op2(e21, e20)) | ~ spl102_298), inference(backward_demodulation, [], [f2609, f2806])).
fof(f2609, plain, ~ (op2(e21, e20) = h2(e11)), inference(backward_demodulation, [], [f404, f909])).
fof(f404, plain, ~ (op2(e21, e20) = op2(e21, e21)), inference(cnf_transformation, [], [f6])).
fof(f5969, plain, (~ spl102_73 | ~ spl102_298), inference(avatar_split_clause, [], [f5962, f2805, f1363])).
fof(f5962, plain, (~ (e20 = op2(e23, e21)) | ~ spl102_298), inference(backward_demodulation, [], [f2608, f2806])).
fof(f2608, plain, ~ (op2(e23, e21) = h2(e11)), inference(backward_demodulation, [], [f384, f909])).
fof(f384, plain, ~ (op2(e21, e21) = op2(e23, e21)), inference(cnf_transformation, [], [f6])).
fof(f5924, plain, (~ spl102_72 | ~ spl102_326), inference(avatar_split_clause, [], [f5915, f2964, f1358])).
fof(f1358, plain, (spl102_72 <=> (e23 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl102_72])])).
fof(f2964, plain, (spl102_326 <=> (e23 = h4(e11))), introduced(avatar_definition, [new_symbols(naming, [spl102_326])])).
fof(f5915, plain, (~ (e23 = op2(e23, e22)) | ~ spl102_326), inference(backward_demodulation, [], [f2626, f2965])).
fof(f2965, plain, ((e23 = h4(e11)) | ~ spl102_326), inference(avatar_component_clause, [], [f2964])).
fof(f2626, plain, ~ (op2(e23, e22) = h4(e11)), inference(backward_demodulation, [], [f421, f917])).
fof(f917, plain, (op2(e23, e23) = h4(e11)), inference(cnf_transformation, [], [f17])).
fof(f17, plain, ((op2(e23, op2(e23, e23)) = h4(e13)) & (h4(e12) = op2(op2(e23, op2(e23, e23)), e23)) & (op2(e23, e23) = h4(e11)) & (e23 = h4(e10))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG109+1.p', ax17)).
fof(f421, plain, ~ (op2(e23, e22) = op2(e23, e23)), inference(cnf_transformation, [], [f6])).
fof(f5923, plain, (~ spl102_100 | ~ spl102_326), inference(avatar_split_clause, [], [f5912, f2964, f1477])).
fof(f5912, plain, (~ (e23 = op2(e21, e23)) | ~ spl102_326), inference(backward_demodulation, [], [f2623, f2965])).
fof(f2623, plain, ~ (op2(e21, e23) = h4(e11)), inference(backward_demodulation, [], [f396, f917])).
fof(f396, plain, ~ (op2(e21, e23) = op2(e23, e23)), inference(cnf_transformation, [], [f6])).
fof(f5900, plain, (~ spl102_72 | ~ spl102_345), inference(avatar_split_clause, [], [f5889, f3085, f1358])).
fof(f3085, plain, (spl102_345 <=> (e23 = h3(e11))), introduced(avatar_definition, [new_symbols(naming, [spl102_345])])).
fof(f5889, plain, (~ (e23 = op2(e23, e22)) | ~ spl102_345), inference(backward_demodulation, [], [f2616, f3086])).
fof(f3086, plain, ((e23 = h3(e11)) | ~ spl102_345), inference(avatar_component_clause, [], [f3085])).
fof(f2616, plain, ~ (op2(e23, e22) = h3(e11)), inference(backward_demodulation, [], [f391, f913])).
fof(f391, plain, ~ (op2(e22, e22) = op2(e23, e22)), inference(cnf_transformation, [], [f6])).
fof(f5826, plain, (~ spl102_82 | ~ spl102_90), inference(avatar_split_clause, [], [f4891, f1435, f1401])).
fof(f1401, plain, (spl102_82 <=> (e21 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl102_82])])).
fof(f1435, plain, (spl102_90 <=> (e21 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl102_90])])).
fof(f4891, plain, (~ (e21 = op2(e22, e23)) | ~ spl102_90), inference(forward_demodulation, [], [f414, f1437])).
fof(f1437, plain, ((e21 = op2(e22, e21)) | ~ spl102_90), inference(avatar_component_clause, [], [f1435])).
fof(f414, plain, ~ (op2(e22, e21) = op2(e22, e23)), inference(cnf_transformation, [], [f6])).
fof(f5821, plain, (~ spl102_75 | ~ spl102_300), inference(avatar_split_clause, [], [f3834, f2816, f1371])).
fof(f1371, plain, (spl102_75 <=> (e22 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl102_75])])).
fof(f2816, plain, (spl102_300 <=> (e22 = h1(e12))), introduced(avatar_definition, [new_symbols(naming, [spl102_300])])).
fof(f3834, plain, (~ (e22 = op2(e23, e21)) | ~ spl102_300), inference(backward_demodulation, [], [f2595, f2817])).
fof(f2817, plain, ((e22 = h1(e12)) | ~ spl102_300), inference(avatar_component_clause, [], [f2816])).
fof(f2595, plain, ~ (op2(e23, e21) = h1(e12)), inference(backward_demodulation, [], [f416, f2591])).
fof(f2591, plain, (op2(e23, e20) = h1(e12)), inference(forward_demodulation, [], [f906, f903])).
fof(f903, plain, (e23 = op2(e20, op2(e20, e20))), inference(cnf_transformation, [], [f13])).
fof(f13, plain, ((e23 = op2(e20, op2(e20, e20))) & (e22 = op2(op2(e20, op2(e20, e20)), e20)) & (op2(e20, e20) = e21)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG109+1.p', ax13)).
fof(f906, plain, (op2(op2(e20, op2(e20, e20)), e20) = h1(e12)), inference(cnf_transformation, [], [f14])).
fof(f14, plain, ((op2(e20, op2(e20, e20)) = h1(e13)) & (op2(op2(e20, op2(e20, e20)), e20) = h1(e12)) & (op2(e20, e20) = h1(e11)) & (e20 = h1(e10))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG109+1.p', ax14)).
fof(f416, plain, ~ (op2(e23, e20) = op2(e23, e21)), inference(cnf_transformation, [], [f6])).
fof(f5819, plain, (~ spl102_74 | ~ spl102_90), inference(avatar_split_clause, [], [f5818, f1435, f1367])).
fof(f1367, plain, (spl102_74 <=> (e21 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl102_74])])).
fof(f5818, plain, (~ (e21 = op2(e23, e21)) | ~ spl102_90), inference(forward_demodulation, [], [f385, f1437])).
fof(f385, plain, ~ (op2(e22, e21) = op2(e23, e21)), inference(cnf_transformation, [], [f6])).
fof(f5816, plain, (~ spl102_20 | ~ spl102_32), inference(avatar_split_clause, [], [f5815, f1156, f1105])).
fof(f1105, plain, (spl102_20 <=> (e13 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl102_20])])).
fof(f5815, plain, (~ (e13 = op1(e12, e13)) | ~ spl102_32), inference(forward_demodulation, [], [f364, f1158])).
fof(f364, plain, ~ (op1(e12, e10) = op1(e12, e13)), inference(cnf_transformation, [], [f5])).
fof(f5814, plain, (~ spl102_7 | ~ spl102_15), inference(avatar_split_clause, [], [f4677, f1084, f1050])).
fof(f1050, plain, (spl102_7 <=> (e12 = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl102_7])])).
fof(f4677, plain, (~ (e12 = op1(e13, e12)) | ~ spl102_15), inference(forward_demodulation, [], [f369, f1086])).
fof(f369, plain, ~ (op1(e13, e10) = op1(e13, e12)), inference(cnf_transformation, [], [f5])).
fof(f5753, plain, (~ spl102_40 | ~ spl102_102 | ~ spl102_300 | ~ spl102_304 | ~ spl102_366), inference(avatar_contradiction_clause, [], [f5752])).
fof(f5752, plain, ($false | (~ spl102_40 | ~ spl102_102 | ~ spl102_300 | ~ spl102_304 | ~ spl102_366)), inference(subsumption_resolution, [], [f5751, f432])).
fof(f432, plain, ~ (e21 = e23), inference(cnf_transformation, [], [f8])).
fof(f5751, plain, ((e21 = e23) | (~ spl102_40 | ~ spl102_102 | ~ spl102_300 | ~ spl102_304 | ~ spl102_366)), inference(forward_demodulation, [], [f2590, f5684])).
fof(f5684, plain, ((e21 = h1(e13)) | (~ spl102_40 | ~ spl102_102 | ~ spl102_300 | ~ spl102_304 | ~ spl102_366)), inference(backward_demodulation, [], [f5629, f1192])).
fof(f5629, plain, ((e21 = h1(op1(e11, e12))) | (~ spl102_102 | ~ spl102_300 | ~ spl102_304 | ~ spl102_366)), inference(forward_demodulation, [], [f5409, f1488])).
fof(f1488, plain, ((e21 = op2(e21, e22)) | ~ spl102_102), inference(avatar_component_clause, [], [f1486])).
fof(f1486, plain, (spl102_102 <=> (e21 = op2(e21, e22))), introduced(avatar_definition, [new_symbols(naming, [spl102_102])])).
fof(f5409, plain, ((op2(e21, e22) = h1(op1(e11, e12))) | (~ spl102_300 | ~ spl102_304 | ~ spl102_366)), inference(forward_demodulation, [], [f5408, f2838])).
fof(f2838, plain, ((e21 = h1(e11)) | ~ spl102_304), inference(avatar_component_clause, [], [f2837])).
fof(f2837, plain, (spl102_304 <=> (e21 = h1(e11))), introduced(avatar_definition, [new_symbols(naming, [spl102_304])])).
fof(f5408, plain, ((h1(op1(e11, e12)) = op2(h1(e11), e22)) | (~ spl102_300 | ~ spl102_366)), inference(forward_demodulation, [], [f3225, f2817])).
fof(f3225, plain, ((h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ spl102_366), inference(avatar_component_clause, [], [f3224])).
fof(f3224, plain, (spl102_366 <=> (h1(op1(e11, e12)) = op2(h1(e11), h1(e12)))), introduced(avatar_definition, [new_symbols(naming, [spl102_366])])).
fof(f2590, plain, (e23 = h1(e13)), inference(forward_demodulation, [], [f907, f903])).
fof(f907, plain, (op2(e20, op2(e20, e20)) = h1(e13)), inference(cnf_transformation, [], [f14])).
fof(f5711, plain, (~ spl102_24 | ~ spl102_32), inference(avatar_split_clause, [], [f5707, f1156, f1122])).
fof(f1122, plain, (spl102_24 <=> (e13 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl102_24])])).
fof(f5707, plain, (~ (e13 = op1(e12, e12)) | ~ spl102_32), inference(backward_demodulation, [], [f363, f1158])).
fof(f363, plain, ~ (op1(e12, e10) = op1(e12, e12)), inference(cnf_transformation, [], [f5])).
fof(f5678, plain, (~ spl102_45 | spl102_61 | ~ spl102_170), inference(avatar_contradiction_clause, [], [f5677])).
fof(f5677, plain, ($false | (~ spl102_45 | spl102_61 | ~ spl102_170)), inference(subsumption_resolution, [], [f5672, f1281])).
fof(f1281, plain, (~ (e10 = op1(e10, e10)) | spl102_61), inference(avatar_component_clause, [], [f1280])).
fof(f1280, plain, (spl102_61 <=> (e10 = op1(e10, e10))), introduced(avatar_definition, [new_symbols(naming, [spl102_61])])).
fof(f5672, plain, ((e10 = op1(e10, e10)) | (~ spl102_45 | ~ spl102_170)), inference(backward_demodulation, [], [f1894, f1214])).
fof(f1214, plain, ((e10 = op1(e11, e10)) | ~ spl102_45), inference(avatar_component_clause, [], [f1212])).
fof(f1212, plain, (spl102_45 <=> (e10 = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl102_45])])).
fof(f1894, plain, ((e10 = op1(e10, op1(e11, e10))) | ~ spl102_170), inference(avatar_component_clause, [], [f1892])).
fof(f1892, plain, (spl102_170 <=> (e10 = op1(e10, op1(e11, e10)))), introduced(avatar_definition, [new_symbols(naming, [spl102_170])])).
fof(f5676, plain, (~ spl102_33 | ~ spl102_45), inference(avatar_split_clause, [], [f5668, f1212, f1161])).
fof(f1161, plain, (spl102_33 <=> (e10 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl102_33])])).
fof(f5668, plain, (~ (e10 = op1(e11, e13)) | ~ spl102_45), inference(backward_demodulation, [], [f358, f1214])).
fof(f358, plain, ~ (op1(e11, e10) = op1(e11, e13)), inference(cnf_transformation, [], [f5])).
fof(f5675, plain, (~ spl102_29 | ~ spl102_45), inference(avatar_split_clause, [], [f5666, f1212, f1144])).
fof(f1144, plain, (spl102_29 <=> (e10 = op1(e12, e10))), introduced(avatar_definition, [new_symbols(naming, [spl102_29])])).
fof(f5666, plain, (~ (e10 = op1(e12, e10)) | ~ spl102_45), inference(backward_demodulation, [], [f329, f1214])).
fof(f329, plain, ~ (op1(e11, e10) = op1(e12, e10)), inference(cnf_transformation, [], [f5])).
fof(f5628, plain, (~ spl102_39 | ~ spl102_55), inference(avatar_split_clause, [], [f5627, f1254, f1186])).
fof(f1186, plain, (spl102_39 <=> (e12 = op1(e11, e12))), introduced(avatar_definition, [new_symbols(naming, [spl102_39])])).
fof(f5627, plain, (~ (e12 = op1(e11, e12)) | ~ spl102_55), inference(forward_demodulation, [], [f338, f1256])).
fof(f338, plain, ~ (op1(e10, e12) = op1(e11, e12)), inference(cnf_transformation, [], [f5])).
fof(f5624, plain, (~ spl102_31 | ~ spl102_15), inference(avatar_split_clause, [], [f4070, f1084, f1152])).
fof(f1152, plain, (spl102_31 <=> (e12 = op1(e12, e10))), introduced(avatar_definition, [new_symbols(naming, [spl102_31])])).
fof(f4070, plain, (~ (e12 = op1(e12, e10)) | ~ spl102_15), inference(forward_demodulation, [], [f331, f1086])).
fof(f331, plain, ~ (op1(e12, e10) = op1(e13, e10)), inference(cnf_transformation, [], [f5])).
fof(f5605, plain, (~ spl102_2 | ~ spl102_4), inference(avatar_contradiction_clause, [], [f5604])).
fof(f5604, plain, ($false | (~ spl102_2 | ~ spl102_4)), inference(subsumption_resolution, [], [f5603, f426])).
fof(f426, plain, ~ (e11 = e13), inference(cnf_transformation, [], [f7])).
fof(f7, plain, (~ (e12 = e13) & ~ (e11 = e13) & ~ (e11 = e12) & ~ (e10 = e13) & ~ (e10 = e12) & ~ (e10 = e11)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG109+1.p', ax7)).
fof(f5603, plain, ((e11 = e13) | (~ spl102_2 | ~ spl102_4)), inference(forward_demodulation, [], [f1039, f1031])).
fof(f1031, plain, ((e11 = op1(e13, e13)) | ~ spl102_2), inference(avatar_component_clause, [], [f1029])).
fof(f1039, plain, ((e13 = op1(e13, e13)) | ~ spl102_4), inference(avatar_component_clause, [], [f1037])).
fof(f5588, plain, (~ spl102_19 | ~ spl102_96 | ~ spl102_265 | ~ spl102_272 | ~ spl102_313), inference(avatar_contradiction_clause, [], [f5587])).
fof(f5587, plain, ($false | (~ spl102_19 | ~ spl102_96 | ~ spl102_265 | ~ spl102_272 | ~ spl102_313)), inference(subsumption_resolution, [], [f5586, f433])).
fof(f433, plain, ~ (e22 = e23), inference(cnf_transformation, [], [f8])).
fof(f5586, plain, ((e22 = e23) | (~ spl102_19 | ~ spl102_96 | ~ spl102_265 | ~ spl102_272 | ~ spl102_313)), inference(forward_demodulation, [], [f5584, f2639])).
fof(f2639, plain, ((e22 = h4(e12)) | ~ spl102_265), inference(avatar_component_clause, [], [f2638])).
fof(f2638, plain, (spl102_265 <=> (e22 = h4(e12))), introduced(avatar_definition, [new_symbols(naming, [spl102_265])])).
fof(f5584, plain, ((e23 = h4(e12)) | (~ spl102_19 | ~ spl102_96 | ~ spl102_265 | ~ spl102_272 | ~ spl102_313)), inference(backward_demodulation, [], [f5450, f1103])).
fof(f1103, plain, ((e12 = op1(e12, e13)) | ~ spl102_19), inference(avatar_component_clause, [], [f1101])).
fof(f5450, plain, ((e23 = h4(op1(e12, e13))) | (~ spl102_96 | ~ spl102_265 | ~ spl102_272 | ~ spl102_313)), inference(forward_demodulation, [], [f5284, f1462])).
fof(f1462, plain, ((e23 = op2(e22, e20)) | ~ spl102_96), inference(avatar_component_clause, [], [f1460])).
fof(f1460, plain, (spl102_96 <=> (e23 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl102_96])])).
fof(f5284, plain, ((op2(e22, e20) = h4(op1(e12, e13))) | (~ spl102_265 | ~ spl102_272 | ~ spl102_313)), inference(forward_demodulation, [], [f5283, f2639])).
fof(f5283, plain, ((h4(op1(e12, e13)) = op2(h4(e12), e20)) | (~ spl102_272 | ~ spl102_313)), inference(forward_demodulation, [], [f2891, f2674])).
fof(f2674, plain, ((e20 = h4(e13)) | ~ spl102_272), inference(avatar_component_clause, [], [f2673])).
fof(f2673, plain, (spl102_272 <=> (e20 = h4(e13))), introduced(avatar_definition, [new_symbols(naming, [spl102_272])])).
fof(f2891, plain, ((h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ spl102_313), inference(avatar_component_clause, [], [f2890])).
fof(f2890, plain, (spl102_313 <=> (h4(op1(e12, e13)) = op2(h4(e12), h4(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl102_313])])).
fof(f5575, plain, (~ spl102_8 | ~ spl102_24), inference(avatar_split_clause, [], [f5571, f1122, f1054])).
fof(f5571, plain, (~ (e13 = op1(e13, e12)) | ~ spl102_24), inference(backward_demodulation, [], [f343, f1124])).
fof(f1124, plain, ((e13 = op1(e12, e12)) | ~ spl102_24), inference(avatar_component_clause, [], [f1122])).
fof(f343, plain, ~ (op1(e12, e12) = op1(e13, e12)), inference(cnf_transformation, [], [f5])).
fof(f5570, plain, (~ spl102_29 | ~ spl102_96 | ~ spl102_300 | ~ spl102_374), inference(avatar_contradiction_clause, [], [f5569])).
fof(f5569, plain, ($false | (~ spl102_29 | ~ spl102_96 | ~ spl102_300 | ~ spl102_374)), inference(subsumption_resolution, [], [f5568, f430])).
fof(f430, plain, ~ (e20 = e23), inference(cnf_transformation, [], [f8])).
fof(f5568, plain, ((e20 = e23) | (~ spl102_29 | ~ spl102_96 | ~ spl102_300 | ~ spl102_374)), inference(forward_demodulation, [], [f5566, f904])).
fof(f904, plain, (e20 = h1(e10)), inference(cnf_transformation, [], [f14])).
fof(f5566, plain, ((e23 = h1(e10)) | (~ spl102_29 | ~ spl102_96 | ~ spl102_300 | ~ spl102_374)), inference(backward_demodulation, [], [f5459, f1146])).
fof(f1146, plain, ((e10 = op1(e12, e10)) | ~ spl102_29), inference(avatar_component_clause, [], [f1144])).
fof(f5459, plain, ((e23 = h1(op1(e12, e10))) | (~ spl102_96 | ~ spl102_300 | ~ spl102_374)), inference(forward_demodulation, [], [f5262, f1462])).
fof(f5262, plain, ((op2(e22, e20) = h1(op1(e12, e10))) | (~ spl102_300 | ~ spl102_374)), inference(forward_demodulation, [], [f3257, f2817])).
fof(f3257, plain, ((h1(op1(e12, e10)) = op2(h1(e12), e20)) | ~ spl102_374), inference(avatar_component_clause, [], [f3256])).
fof(f3256, plain, (spl102_374 <=> (h1(op1(e12, e10)) = op2(h1(e12), e20))), introduced(avatar_definition, [new_symbols(naming, [spl102_374])])).
fof(f5567, plain, (~ spl102_17 | ~ spl102_29), inference(avatar_split_clause, [], [f5560, f1144, f1093])).
fof(f5560, plain, (~ (e10 = op1(e12, e13)) | ~ spl102_29), inference(backward_demodulation, [], [f364, f1146])).
fof(f5542, plain, (~ spl102_48 | ~ spl102_109 | ~ spl102_304 | ~ spl102_376), inference(avatar_contradiction_clause, [], [f5541])).
fof(f5541, plain, ($false | (~ spl102_48 | ~ spl102_109 | ~ spl102_304 | ~ spl102_376)), inference(subsumption_resolution, [], [f5540, f430])).
fof(f5540, plain, ((e20 = e23) | (~ spl102_48 | ~ spl102_109 | ~ spl102_304 | ~ spl102_376)), inference(backward_demodulation, [], [f2590, f5536])).
fof(f5536, plain, ((e20 = h1(e13)) | (~ spl102_48 | ~ spl102_109 | ~ spl102_304 | ~ spl102_376)), inference(backward_demodulation, [], [f5468, f1226])).
fof(f1226, plain, ((e13 = op1(e11, e10)) | ~ spl102_48), inference(avatar_component_clause, [], [f1224])).
fof(f1224, plain, (spl102_48 <=> (e13 = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl102_48])])).
fof(f5468, plain, ((e20 = h1(op1(e11, e10))) | (~ spl102_109 | ~ spl102_304 | ~ spl102_376)), inference(forward_demodulation, [], [f5232, f1518])).
fof(f5232, plain, ((op2(e21, e20) = h1(op1(e11, e10))) | (~ spl102_304 | ~ spl102_376)), inference(forward_demodulation, [], [f3265, f2838])).
fof(f3265, plain, ((h1(op1(e11, e10)) = op2(h1(e11), e20)) | ~ spl102_376), inference(avatar_component_clause, [], [f3264])).
fof(f3264, plain, (spl102_376 <=> (h1(op1(e11, e10)) = op2(h1(e11), e20))), introduced(avatar_definition, [new_symbols(naming, [spl102_376])])).
fof(f5538, plain, (~ spl102_32 | ~ spl102_48), inference(avatar_split_clause, [], [f5533, f1224, f1156])).
fof(f5533, plain, (~ (e13 = op1(e12, e10)) | ~ spl102_48), inference(backward_demodulation, [], [f329, f1226])).
fof(f5527, plain, (~ spl102_33 | ~ spl102_49), inference(avatar_split_clause, [], [f5519, f1229, f1161])).
fof(f5519, plain, (~ (e10 = op1(e11, e13)) | ~ spl102_49), inference(backward_demodulation, [], [f344, f1231])).
fof(f344, plain, ~ (op1(e10, e13) = op1(e11, e13)), inference(cnf_transformation, [], [f5])).
fof(f5513, plain, (~ spl102_55 | ~ spl102_102 | ~ spl102_289 | spl102_293 | ~ spl102_360), inference(avatar_contradiction_clause, [], [f5512])).
fof(f5512, plain, ($false | (~ spl102_55 | ~ spl102_102 | ~ spl102_289 | spl102_293 | ~ spl102_360)), inference(subsumption_resolution, [], [f5506, f2781])).
fof(f2781, plain, (~ (e21 = h2(e12)) | spl102_293), inference(avatar_component_clause, [], [f2779])).
fof(f2779, plain, (spl102_293 <=> (e21 = h2(e12))), introduced(avatar_definition, [new_symbols(naming, [spl102_293])])).
fof(f5506, plain, ((e21 = h2(e12)) | (~ spl102_55 | ~ spl102_102 | ~ spl102_289 | ~ spl102_360)), inference(backward_demodulation, [], [f5476, f1256])).
fof(f5476, plain, ((e21 = h2(op1(e10, e12))) | (~ spl102_102 | ~ spl102_289 | ~ spl102_360)), inference(forward_demodulation, [], [f5269, f1488])).
fof(f5269, plain, ((op2(e21, e22) = h2(op1(e10, e12))) | (~ spl102_289 | ~ spl102_360)), inference(forward_demodulation, [], [f3158, f2760])).
fof(f2760, plain, ((e22 = h2(e12)) | ~ spl102_289), inference(avatar_component_clause, [], [f2759])).
fof(f2759, plain, (spl102_289 <=> (e22 = h2(e12))), introduced(avatar_definition, [new_symbols(naming, [spl102_289])])).
fof(f3158, plain, ((h2(op1(e10, e12)) = op2(e21, h2(e12))) | ~ spl102_360), inference(avatar_component_clause, [], [f3157])).
fof(f3157, plain, (spl102_360 <=> (h2(op1(e10, e12)) = op2(e21, h2(e12)))), introduced(avatar_definition, [new_symbols(naming, [spl102_360])])).
fof(f5494, plain, (spl102_48 | ~ spl102_60 | ~ spl102_171), inference(avatar_split_clause, [], [f5493, f1897, f1275, f1224])).
fof(f1275, plain, (spl102_60 <=> (e13 = op1(e10, e11))), introduced(avatar_definition, [new_symbols(naming, [spl102_60])])).
fof(f1897, plain, (spl102_171 <=> (op1(e10, e11) = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl102_171])])).
fof(f5493, plain, ((e13 = op1(e11, e10)) | (~ spl102_60 | ~ spl102_171)), inference(forward_demodulation, [], [f1899, f1277])).
fof(f1277, plain, ((e13 = op1(e10, e11)) | ~ spl102_60), inference(avatar_component_clause, [], [f1275])).
fof(f1899, plain, ((op1(e10, e11) = op1(e11, e10)) | ~ spl102_171), inference(avatar_component_clause, [], [f1897])).
fof(f5466, plain, (~ spl102_33 | ~ spl102_9 | spl102_149), inference(avatar_split_clause, [], [f5465, f1791, f1059, f1161])).
fof(f1059, plain, (spl102_9 <=> (e10 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl102_9])])).
fof(f5465, plain, (~ (e10 = op1(e11, e13)) | (~ spl102_9 | spl102_149)), inference(forward_demodulation, [], [f1792, f1061])).
fof(f1061, plain, ((e10 = op1(e13, e11)) | ~ spl102_9), inference(avatar_component_clause, [], [f1059])).
fof(f1792, plain, (~ (op1(e11, e13) = op1(e13, e11)) | spl102_149), inference(avatar_component_clause, [], [f1791])).
fof(f5443, plain, (~ spl102_5 | ~ spl102_9), inference(avatar_split_clause, [], [f4871, f1059, f1042])).
fof(f4871, plain, (~ (e10 = op1(e13, e12)) | ~ spl102_9), inference(backward_demodulation, [], [f371, f1061])).
fof(f371, plain, ~ (op1(e13, e11) = op1(e13, e12)), inference(cnf_transformation, [], [f5])).
fof(f5442, plain, (~ spl102_90 | ~ spl102_102 | spl102_226), inference(avatar_contradiction_clause, [], [f5441])).
fof(f5441, plain, ($false | (~ spl102_90 | ~ spl102_102 | spl102_226)), inference(subsumption_resolution, [], [f5440, f1488])).
fof(f5440, plain, (~ (e21 = op2(e21, e22)) | (~ spl102_90 | spl102_226)), inference(forward_demodulation, [], [f2316, f1437])).
fof(f2316, plain, (~ (op2(e21, e22) = op2(e22, e21)) | spl102_226), inference(avatar_component_clause, [], [f2315])).
fof(f2315, plain, (spl102_226 <=> (op2(e21, e22) = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl102_226])])).
fof(f5439, plain, (~ spl102_38 | ~ spl102_102 | ~ spl102_265 | ~ spl102_270 | spl102_309), inference(avatar_contradiction_clause, [], [f5438])).
fof(f5438, plain, ($false | (~ spl102_38 | ~ spl102_102 | ~ spl102_265 | ~ spl102_270 | spl102_309)), inference(subsumption_resolution, [], [f5437, f1488])).
fof(f5437, plain, (~ (e21 = op2(e21, e22)) | (~ spl102_38 | ~ spl102_265 | ~ spl102_270 | spl102_309)), inference(forward_demodulation, [], [f5436, f2664])).
fof(f2664, plain, ((e21 = h4(e11)) | ~ spl102_270), inference(avatar_component_clause, [], [f2663])).
fof(f2663, plain, (spl102_270 <=> (e21 = h4(e11))), introduced(avatar_definition, [new_symbols(naming, [spl102_270])])).
fof(f5436, plain, (~ (op2(e21, e22) = h4(e11)) | (~ spl102_38 | ~ spl102_265 | ~ spl102_270 | spl102_309)), inference(forward_demodulation, [], [f5435, f1184])).
fof(f5435, plain, (~ (op2(e21, e22) = h4(op1(e11, e12))) | (~ spl102_265 | ~ spl102_270 | spl102_309)), inference(forward_demodulation, [], [f5434, f2664])).
fof(f5434, plain, (~ (h4(op1(e11, e12)) = op2(h4(e11), e22)) | (~ spl102_265 | spl102_309)), inference(forward_demodulation, [], [f2876, f2639])).
fof(f2876, plain, (~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | spl102_309), inference(avatar_component_clause, [], [f2874])).
fof(f2874, plain, (spl102_309 <=> (h4(op1(e11, e12)) = op2(h4(e11), h4(e12)))), introduced(avatar_definition, [new_symbols(naming, [spl102_309])])).
fof(f5419, plain, (~ spl102_62 | spl102_324), inference(avatar_contradiction_clause, [], [f5418])).
fof(f5418, plain, ($false | (~ spl102_62 | spl102_324)), inference(trivial_inequality_removal, [], [f5417])).
fof(f5417, plain, (~ (h4(e11) = h4(e11)) | (~ spl102_62 | spl102_324)), inference(forward_demodulation, [], [f2936, f1286])).
fof(f1286, plain, ((op1(e10, e10) = e11) | ~ spl102_62), inference(avatar_component_clause, [], [f1284])).
fof(f1284, plain, (spl102_62 <=> (op1(e10, e10) = e11)), introduced(avatar_definition, [new_symbols(naming, [spl102_62])])).
fof(f2936, plain, (~ (h4(e11) = h4(op1(e10, e10))) | spl102_324), inference(avatar_component_clause, [], [f2934])).
fof(f2934, plain, (spl102_324 <=> (h4(e11) = h4(op1(e10, e10)))), introduced(avatar_definition, [new_symbols(naming, [spl102_324])])).
fof(f5404, plain, (~ spl102_51 | ~ spl102_79 | ~ spl102_265 | ~ spl102_272 | spl102_321), inference(avatar_contradiction_clause, [], [f5403])).
fof(f5403, plain, ($false | (~ spl102_51 | ~ spl102_79 | ~ spl102_265 | ~ spl102_272 | spl102_321)), inference(subsumption_resolution, [], [f5402, f1390])).
fof(f5402, plain, (~ (e22 = op2(e23, e20)) | (~ spl102_51 | ~ spl102_265 | ~ spl102_272 | spl102_321)), inference(forward_demodulation, [], [f5401, f2639])).
fof(f5401, plain, (~ (op2(e23, e20) = h4(e12)) | (~ spl102_51 | ~ spl102_272 | spl102_321)), inference(forward_demodulation, [], [f5400, f1239])).
fof(f5400, plain, (~ (op2(e23, e20) = h4(op1(e10, e13))) | (~ spl102_272 | spl102_321)), inference(forward_demodulation, [], [f2924, f2674])).
fof(f2924, plain, (~ (h4(op1(e10, e13)) = op2(e23, h4(e13))) | spl102_321), inference(avatar_component_clause, [], [f2922])).
fof(f2922, plain, (spl102_321 <=> (h4(op1(e10, e13)) = op2(e23, h4(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl102_321])])).
fof(f5395, plain, (~ spl102_38 | ~ spl102_102 | ~ spl102_300 | ~ spl102_304 | spl102_366), inference(avatar_contradiction_clause, [], [f5394])).
fof(f5394, plain, ($false | (~ spl102_38 | ~ spl102_102 | ~ spl102_300 | ~ spl102_304 | spl102_366)), inference(subsumption_resolution, [], [f5393, f1488])).
fof(f5393, plain, (~ (e21 = op2(e21, e22)) | (~ spl102_38 | ~ spl102_300 | ~ spl102_304 | spl102_366)), inference(forward_demodulation, [], [f5392, f2838])).
fof(f5392, plain, (~ (op2(e21, e22) = h1(e11)) | (~ spl102_38 | ~ spl102_300 | ~ spl102_304 | spl102_366)), inference(forward_demodulation, [], [f5391, f1184])).
fof(f5391, plain, (~ (op2(e21, e22) = h1(op1(e11, e12))) | (~ spl102_300 | ~ spl102_304 | spl102_366)), inference(forward_demodulation, [], [f5390, f2838])).
fof(f5390, plain, (~ (h1(op1(e11, e12)) = op2(h1(e11), e22)) | (~ spl102_300 | spl102_366)), inference(forward_demodulation, [], [f3226, f2817])).
fof(f3226, plain, (~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | spl102_366), inference(avatar_component_clause, [], [f3224])).
fof(f5383, plain, (~ spl102_45 | ~ spl102_100 | ~ spl102_270 | spl102_320), inference(avatar_contradiction_clause, [], [f5382])).
fof(f5382, plain, ($false | (~ spl102_45 | ~ spl102_100 | ~ spl102_270 | spl102_320)), inference(subsumption_resolution, [], [f5381, f1479])).
fof(f5381, plain, (~ (e23 = op2(e21, e23)) | (~ spl102_45 | ~ spl102_270 | spl102_320)), inference(forward_demodulation, [], [f5380, f916])).
fof(f916, plain, (e23 = h4(e10)), inference(cnf_transformation, [], [f17])).
fof(f5380, plain, (~ (op2(e21, e23) = h4(e10)) | (~ spl102_45 | ~ spl102_270 | spl102_320)), inference(forward_demodulation, [], [f5379, f1214])).
fof(f5379, plain, (~ (op2(e21, e23) = h4(op1(e11, e10))) | (~ spl102_270 | spl102_320)), inference(forward_demodulation, [], [f2920, f2664])).
fof(f2920, plain, (~ (h4(op1(e11, e10)) = op2(h4(e11), e23)) | spl102_320), inference(avatar_component_clause, [], [f2918])).
fof(f2918, plain, (spl102_320 <=> (h4(op1(e11, e10)) = op2(h4(e11), e23))), introduced(avatar_definition, [new_symbols(naming, [spl102_320])])).
fof(f5359, plain, (~ spl102_32 | ~ spl102_81 | ~ spl102_265 | ~ spl102_272 | spl102_319), inference(avatar_contradiction_clause, [], [f5358])).
fof(f5358, plain, ($false | (~ spl102_32 | ~ spl102_81 | ~ spl102_265 | ~ spl102_272 | spl102_319)), inference(subsumption_resolution, [], [f5357, f1399])).
fof(f1399, plain, ((e20 = op2(e22, e23)) | ~ spl102_81), inference(avatar_component_clause, [], [f1397])).
fof(f1397, plain, (spl102_81 <=> (e20 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl102_81])])).
fof(f5357, plain, (~ (e20 = op2(e22, e23)) | (~ spl102_32 | ~ spl102_265 | ~ spl102_272 | spl102_319)), inference(forward_demodulation, [], [f5356, f2674])).
fof(f5356, plain, (~ (op2(e22, e23) = h4(e13)) | (~ spl102_32 | ~ spl102_265 | spl102_319)), inference(forward_demodulation, [], [f5355, f1158])).
fof(f5355, plain, (~ (op2(e22, e23) = h4(op1(e12, e10))) | (~ spl102_265 | spl102_319)), inference(forward_demodulation, [], [f2916, f2639])).
fof(f2916, plain, (~ (h4(op1(e12, e10)) = op2(h4(e12), e23)) | spl102_319), inference(avatar_component_clause, [], [f2914])).
fof(f2914, plain, (spl102_319 <=> (h4(op1(e12, e10)) = op2(h4(e12), e23))), introduced(avatar_definition, [new_symbols(naming, [spl102_319])])).
fof(f5333, plain, (~ spl102_2 | ~ spl102_126 | ~ spl102_270 | ~ spl102_272 | spl102_316), inference(avatar_contradiction_clause, [], [f5332])).
fof(f5332, plain, ($false | (~ spl102_2 | ~ spl102_126 | ~ spl102_270 | ~ spl102_272 | spl102_316)), inference(subsumption_resolution, [], [f5331, f1590])).
fof(f1590, plain, ((op2(e20, e20) = e21) | ~ spl102_126), inference(avatar_component_clause, [], [f1588])).
fof(f1588, plain, (spl102_126 <=> (op2(e20, e20) = e21)), introduced(avatar_definition, [new_symbols(naming, [spl102_126])])).
fof(f5331, plain, (~ (op2(e20, e20) = e21) | (~ spl102_2 | ~ spl102_270 | ~ spl102_272 | spl102_316)), inference(forward_demodulation, [], [f5330, f2664])).
fof(f5330, plain, (~ (op2(e20, e20) = h4(e11)) | (~ spl102_2 | ~ spl102_272 | spl102_316)), inference(forward_demodulation, [], [f5329, f1031])).
fof(f5329, plain, (~ (op2(e20, e20) = h4(op1(e13, e13))) | (~ spl102_272 | spl102_316)), inference(forward_demodulation, [], [f2904, f2674])).
fof(f2904, plain, (~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | spl102_316), inference(avatar_component_clause, [], [f2902])).
fof(f2902, plain, (spl102_316 <=> (h4(op1(e13, e13)) = op2(h4(e13), h4(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl102_316])])).
fof(f5306, plain, (~ spl102_8 | ~ spl102_117 | ~ spl102_265 | ~ spl102_272 | spl102_315), inference(avatar_contradiction_clause, [], [f5305])).
fof(f5305, plain, ($false | (~ spl102_8 | ~ spl102_117 | ~ spl102_265 | ~ spl102_272 | spl102_315)), inference(subsumption_resolution, [], [f5304, f1552])).
fof(f5304, plain, (~ (e20 = op2(e20, e22)) | (~ spl102_8 | ~ spl102_265 | ~ spl102_272 | spl102_315)), inference(forward_demodulation, [], [f5303, f2674])).
fof(f5303, plain, (~ (op2(e20, e22) = h4(e13)) | (~ spl102_8 | ~ spl102_265 | ~ spl102_272 | spl102_315)), inference(forward_demodulation, [], [f5302, f1056])).
fof(f1056, plain, ((e13 = op1(e13, e12)) | ~ spl102_8), inference(avatar_component_clause, [], [f1054])).
fof(f5302, plain, (~ (op2(e20, e22) = h4(op1(e13, e12))) | (~ spl102_265 | ~ spl102_272 | spl102_315)), inference(forward_demodulation, [], [f5301, f2674])).
fof(f5301, plain, (~ (h4(op1(e13, e12)) = op2(h4(e13), e22)) | (~ spl102_265 | spl102_315)), inference(forward_demodulation, [], [f2900, f2639])).
fof(f2900, plain, (~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | spl102_315), inference(avatar_component_clause, [], [f2898])).
fof(f2898, plain, (spl102_315 <=> (h4(op1(e13, e12)) = op2(h4(e13), h4(e12)))), introduced(avatar_definition, [new_symbols(naming, [spl102_315])])).
fof(f5282, plain, (~ spl102_9 | ~ spl102_124 | ~ spl102_270 | ~ spl102_272 | spl102_314), inference(avatar_contradiction_clause, [], [f5281])).
fof(f5281, plain, ($false | (~ spl102_9 | ~ spl102_124 | ~ spl102_270 | ~ spl102_272 | spl102_314)), inference(subsumption_resolution, [], [f5280, f1581])).
fof(f5280, plain, (~ (e23 = op2(e20, e21)) | (~ spl102_9 | ~ spl102_270 | ~ spl102_272 | spl102_314)), inference(forward_demodulation, [], [f5279, f916])).
fof(f5279, plain, (~ (op2(e20, e21) = h4(e10)) | (~ spl102_9 | ~ spl102_270 | ~ spl102_272 | spl102_314)), inference(forward_demodulation, [], [f5278, f1061])).
fof(f5278, plain, (~ (op2(e20, e21) = h4(op1(e13, e11))) | (~ spl102_270 | ~ spl102_272 | spl102_314)), inference(forward_demodulation, [], [f5277, f2674])).
fof(f5277, plain, (~ (h4(op1(e13, e11)) = op2(h4(e13), e21)) | (~ spl102_270 | spl102_314)), inference(forward_demodulation, [], [f2896, f2664])).
fof(f2896, plain, (~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | spl102_314), inference(avatar_component_clause, [], [f2894])).
fof(f2894, plain, (spl102_314 <=> (h4(op1(e13, e11)) = op2(h4(e13), h4(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl102_314])])).
fof(f5257, plain, (~ spl102_17 | ~ spl102_96 | ~ spl102_265 | ~ spl102_272 | spl102_313), inference(avatar_contradiction_clause, [], [f5256])).
fof(f5256, plain, ($false | (~ spl102_17 | ~ spl102_96 | ~ spl102_265 | ~ spl102_272 | spl102_313)), inference(subsumption_resolution, [], [f5255, f1462])).
fof(f5255, plain, (~ (e23 = op2(e22, e20)) | (~ spl102_17 | ~ spl102_265 | ~ spl102_272 | spl102_313)), inference(forward_demodulation, [], [f5254, f916])).
fof(f5254, plain, (~ (op2(e22, e20) = h4(e10)) | (~ spl102_17 | ~ spl102_265 | ~ spl102_272 | spl102_313)), inference(forward_demodulation, [], [f5253, f1095])).
fof(f1095, plain, ((e10 = op1(e12, e13)) | ~ spl102_17), inference(avatar_component_clause, [], [f1093])).
fof(f5253, plain, (~ (op2(e22, e20) = h4(op1(e12, e13))) | (~ spl102_265 | ~ spl102_272 | spl102_313)), inference(forward_demodulation, [], [f5252, f2639])).
fof(f5252, plain, (~ (h4(op1(e12, e13)) = op2(h4(e12), e20)) | (~ spl102_272 | spl102_313)), inference(forward_demodulation, [], [f2892, f2674])).
fof(f2892, plain, (~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | spl102_313), inference(avatar_component_clause, [], [f2890])).
fof(f5244, plain, (~ spl102_53 | ~ spl102_102 | ~ spl102_289 | spl102_360), inference(avatar_contradiction_clause, [], [f5243])).
fof(f5243, plain, ($false | (~ spl102_53 | ~ spl102_102 | ~ spl102_289 | spl102_360)), inference(subsumption_resolution, [], [f5242, f1488])).
fof(f5242, plain, (~ (e21 = op2(e21, e22)) | (~ spl102_53 | ~ spl102_289 | spl102_360)), inference(forward_demodulation, [], [f5241, f908])).
fof(f908, plain, (e21 = h2(e10)), inference(cnf_transformation, [], [f15])).
fof(f5241, plain, (~ (op2(e21, e22) = h2(e10)) | (~ spl102_53 | ~ spl102_289 | spl102_360)), inference(forward_demodulation, [], [f5240, f1248])).
fof(f5240, plain, (~ (op2(e21, e22) = h2(op1(e10, e12))) | (~ spl102_289 | spl102_360)), inference(forward_demodulation, [], [f3159, f2760])).
fof(f3159, plain, (~ (h2(op1(e10, e12)) = op2(e21, h2(e12))) | spl102_360), inference(avatar_component_clause, [], [f3157])).
fof(f5239, plain, (~ spl102_32 | ~ spl102_96 | ~ spl102_300 | spl102_374), inference(avatar_contradiction_clause, [], [f5238])).
fof(f5238, plain, ($false | (~ spl102_32 | ~ spl102_96 | ~ spl102_300 | spl102_374)), inference(subsumption_resolution, [], [f5237, f1462])).
fof(f5237, plain, (~ (e23 = op2(e22, e20)) | (~ spl102_32 | ~ spl102_300 | spl102_374)), inference(forward_demodulation, [], [f5236, f2590])).
fof(f5236, plain, (~ (op2(e22, e20) = h1(e13)) | (~ spl102_32 | ~ spl102_300 | spl102_374)), inference(forward_demodulation, [], [f5235, f1158])).
fof(f5235, plain, (~ (op2(e22, e20) = h1(op1(e12, e10))) | (~ spl102_300 | spl102_374)), inference(forward_demodulation, [], [f3258, f2817])).
fof(f3258, plain, (~ (h1(op1(e12, e10)) = op2(h1(e12), e20)) | spl102_374), inference(avatar_component_clause, [], [f3256])).
fof(f5227, plain, (~ spl102_26 | ~ spl102_90 | ~ spl102_265 | ~ spl102_270 | spl102_311), inference(avatar_contradiction_clause, [], [f5226])).
fof(f5226, plain, ($false | (~ spl102_26 | ~ spl102_90 | ~ spl102_265 | ~ spl102_270 | spl102_311)), inference(subsumption_resolution, [], [f5225, f1437])).
fof(f5225, plain, (~ (e21 = op2(e22, e21)) | (~ spl102_26 | ~ spl102_265 | ~ spl102_270 | spl102_311)), inference(forward_demodulation, [], [f5224, f2664])).
fof(f5224, plain, (~ (op2(e22, e21) = h4(e11)) | (~ spl102_26 | ~ spl102_265 | ~ spl102_270 | spl102_311)), inference(forward_demodulation, [], [f5223, f1133])).
fof(f1133, plain, ((e11 = op1(e12, e11)) | ~ spl102_26), inference(avatar_component_clause, [], [f1131])).
fof(f5223, plain, (~ (op2(e22, e21) = h4(op1(e12, e11))) | (~ spl102_265 | ~ spl102_270 | spl102_311)), inference(forward_demodulation, [], [f5222, f2639])).
fof(f5222, plain, (~ (h4(op1(e12, e11)) = op2(h4(e12), e21)) | (~ spl102_270 | spl102_311)), inference(forward_demodulation, [], [f2884, f2664])).
fof(f2884, plain, (~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | spl102_311), inference(avatar_component_clause, [], [f2882])).
fof(f2882, plain, (spl102_311 <=> (h4(op1(e12, e11)) = op2(h4(e12), h4(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl102_311])])).
fof(f5204, plain, (~ spl102_45 | ~ spl102_109 | ~ spl102_304 | spl102_376), inference(avatar_contradiction_clause, [], [f5203])).
fof(f5203, plain, ($false | (~ spl102_45 | ~ spl102_109 | ~ spl102_304 | spl102_376)), inference(subsumption_resolution, [], [f5202, f1518])).
fof(f5202, plain, (~ (e20 = op2(e21, e20)) | (~ spl102_45 | ~ spl102_304 | spl102_376)), inference(forward_demodulation, [], [f5201, f904])).
fof(f5201, plain, (~ (op2(e21, e20) = h1(e10)) | (~ spl102_45 | ~ spl102_304 | spl102_376)), inference(forward_demodulation, [], [f5200, f1214])).
fof(f5200, plain, (~ (op2(e21, e20) = h1(op1(e11, e10))) | (~ spl102_304 | spl102_376)), inference(forward_demodulation, [], [f3266, f2838])).
fof(f3266, plain, (~ (h1(op1(e11, e10)) = op2(h1(e11), e20)) | spl102_376), inference(avatar_component_clause, [], [f3264])).
fof(f5197, plain, (~ spl102_36 | ~ spl102_109 | ~ spl102_270 | ~ spl102_272 | spl102_310), inference(avatar_contradiction_clause, [], [f5196])).
fof(f5196, plain, ($false | (~ spl102_36 | ~ spl102_109 | ~ spl102_270 | ~ spl102_272 | spl102_310)), inference(subsumption_resolution, [], [f5195, f1518])).
fof(f5195, plain, (~ (e20 = op2(e21, e20)) | (~ spl102_36 | ~ spl102_270 | ~ spl102_272 | spl102_310)), inference(forward_demodulation, [], [f5194, f2674])).
fof(f5194, plain, (~ (op2(e21, e20) = h4(e13)) | (~ spl102_36 | ~ spl102_270 | ~ spl102_272 | spl102_310)), inference(forward_demodulation, [], [f5193, f1175])).
fof(f5193, plain, (~ (op2(e21, e20) = h4(op1(e11, e13))) | (~ spl102_270 | ~ spl102_272 | spl102_310)), inference(forward_demodulation, [], [f5192, f2664])).
fof(f5192, plain, (~ (h4(op1(e11, e13)) = op2(h4(e11), e20)) | (~ spl102_272 | spl102_310)), inference(forward_demodulation, [], [f2880, f2674])).
fof(f2880, plain, (~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | spl102_310), inference(avatar_component_clause, [], [f2878])).
fof(f2878, plain, (spl102_310 <=> (h4(op1(e11, e13)) = op2(h4(e11), h4(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl102_310])])).
fof(f5155, plain, (~ spl102_90 | spl102_106 | ~ spl102_225), inference(avatar_contradiction_clause, [], [f5154])).
fof(f5154, plain, ($false | (~ spl102_90 | spl102_106 | ~ spl102_225)), inference(subsumption_resolution, [], [f5153, f1504])).
fof(f1504, plain, (~ (e21 = op2(e21, e21)) | spl102_106), inference(avatar_component_clause, [], [f1503])).
fof(f1503, plain, (spl102_106 <=> (e21 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl102_106])])).
fof(f5153, plain, ((e21 = op2(e21, e21)) | (~ spl102_90 | ~ spl102_225)), inference(forward_demodulation, [], [f2312, f1437])).
fof(f2312, plain, ((e21 = op2(e21, op2(e22, e21))) | ~ spl102_225), inference(avatar_component_clause, [], [f2310])).
fof(f2310, plain, (spl102_225 <=> (e21 = op2(e21, op2(e22, e21)))), introduced(avatar_definition, [new_symbols(naming, [spl102_225])])).
fof(f5152, plain, (~ spl102_23 | ~ spl102_87 | ~ spl102_265 | spl102_312), inference(avatar_contradiction_clause, [], [f5151])).
fof(f5151, plain, ($false | (~ spl102_23 | ~ spl102_87 | ~ spl102_265 | spl102_312)), inference(subsumption_resolution, [], [f5148, f1424])).
fof(f1424, plain, ((e22 = op2(e22, e22)) | ~ spl102_87), inference(avatar_component_clause, [], [f1422])).
fof(f1422, plain, (spl102_87 <=> (e22 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl102_87])])).
fof(f5148, plain, (~ (e22 = op2(e22, e22)) | (~ spl102_23 | ~ spl102_265 | spl102_312)), inference(backward_demodulation, [], [f5089, f2639])).
fof(f5089, plain, (~ (h4(e12) = op2(h4(e12), h4(e12))) | (~ spl102_23 | spl102_312)), inference(forward_demodulation, [], [f2888, f1120])).
fof(f2888, plain, (~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | spl102_312), inference(avatar_component_clause, [], [f2886])).
fof(f2886, plain, (spl102_312 <=> (h4(op1(e12, e12)) = op2(h4(e12), h4(e12)))), introduced(avatar_definition, [new_symbols(naming, [spl102_312])])).
fof(f5150, plain, (~ spl102_53 | ~ spl102_72 | ~ spl102_265 | spl102_322), inference(avatar_contradiction_clause, [], [f5149])).
fof(f5149, plain, ($false | (~ spl102_53 | ~ spl102_72 | ~ spl102_265 | spl102_322)), inference(subsumption_resolution, [], [f5147, f1360])).
fof(f1360, plain, ((e23 = op2(e23, e22)) | ~ spl102_72), inference(avatar_component_clause, [], [f1358])).
fof(f5147, plain, (~ (e23 = op2(e23, e22)) | (~ spl102_53 | ~ spl102_265 | spl102_322)), inference(backward_demodulation, [], [f5093, f2639])).
fof(f5093, plain, (~ (e23 = op2(e23, h4(e12))) | (~ spl102_53 | spl102_322)), inference(forward_demodulation, [], [f5092, f916])).
fof(f5092, plain, (~ (h4(e10) = op2(e23, h4(e12))) | (~ spl102_53 | spl102_322)), inference(forward_demodulation, [], [f2928, f1248])).
fof(f2928, plain, (~ (h4(op1(e10, e12)) = op2(e23, h4(e12))) | spl102_322), inference(avatar_component_clause, [], [f2926])).
fof(f2926, plain, (spl102_322 <=> (h4(op1(e10, e12)) = op2(e23, h4(e12)))), introduced(avatar_definition, [new_symbols(naming, [spl102_322])])).
fof(f5091, plain, (spl102_265 | ~ spl102_115 | ~ spl102_272), inference(avatar_split_clause, [], [f5090, f2673, f1541, f2638])).
fof(f5090, plain, ((e22 = h4(e12)) | (~ spl102_115 | ~ spl102_272)), inference(forward_demodulation, [], [f5041, f1543])).
fof(f1543, plain, ((e22 = op2(e20, e23)) | ~ spl102_115), inference(avatar_component_clause, [], [f1541])).
fof(f5041, plain, ((op2(e20, e23) = h4(e12)) | ~ spl102_272), inference(forward_demodulation, [], [f2621, f2674])).
fof(f2621, plain, (h4(e12) = op2(h4(e13), e23)), inference(forward_demodulation, [], [f918, f919])).
fof(f919, plain, (op2(e23, op2(e23, e23)) = h4(e13)), inference(cnf_transformation, [], [f17])).
fof(f918, plain, (h4(e12) = op2(op2(e23, op2(e23, e23)), e23)), inference(cnf_transformation, [], [f17])).
fof(f5088, plain, (~ spl102_265 | ~ spl102_43 | ~ spl102_107 | ~ spl102_270 | spl102_308), inference(avatar_split_clause, [], [f5087, f2870, f2663, f1507, f1203, f2638])).
fof(f1203, plain, (spl102_43 <=> (e12 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl102_43])])).
fof(f1507, plain, (spl102_107 <=> (e22 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl102_107])])).
fof(f2870, plain, (spl102_308 <=> (h4(op1(e11, e11)) = op2(h4(e11), h4(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl102_308])])).
fof(f5087, plain, (~ (e22 = h4(e12)) | (~ spl102_43 | ~ spl102_107 | ~ spl102_270 | spl102_308)), inference(forward_demodulation, [], [f5064, f1509])).
fof(f1509, plain, ((e22 = op2(e21, e21)) | ~ spl102_107), inference(avatar_component_clause, [], [f1507])).
fof(f5064, plain, (~ (op2(e21, e21) = h4(e12)) | (~ spl102_43 | ~ spl102_270 | spl102_308)), inference(forward_demodulation, [], [f5063, f1205])).
fof(f1205, plain, ((e12 = op1(e11, e11)) | ~ spl102_43), inference(avatar_component_clause, [], [f1203])).
fof(f5063, plain, (~ (op2(e21, e21) = h4(op1(e11, e11))) | (~ spl102_270 | spl102_308)), inference(forward_demodulation, [], [f2872, f2664])).
fof(f2872, plain, (~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | spl102_308), inference(avatar_component_clause, [], [f2870])).
fof(f5086, plain, (spl102_289 | ~ spl102_107 | ~ spl102_292), inference(avatar_split_clause, [], [f5085, f2774, f1507, f2759])).
fof(f5085, plain, ((e22 = h2(e12)) | (~ spl102_107 | ~ spl102_292)), inference(forward_demodulation, [], [f5084, f1509])).
fof(f5084, plain, ((op2(e21, e21) = h2(e12)) | ~ spl102_292), inference(forward_demodulation, [], [f2605, f2775])).
fof(f2605, plain, (h2(e12) = op2(h2(e13), e21)), inference(forward_demodulation, [], [f910, f911])).
fof(f910, plain, (h2(e12) = op2(op2(e21, op2(e21, e21)), e21)), inference(cnf_transformation, [], [f15])).
fof(f5079, plain, (~ spl102_90 | ~ spl102_102 | ~ spl102_233), inference(avatar_contradiction_clause, [], [f5078])).
fof(f5078, plain, ($false | (~ spl102_90 | ~ spl102_102 | ~ spl102_233)), inference(subsumption_resolution, [], [f5077, f431])).
fof(f431, plain, ~ (e21 = e22), inference(cnf_transformation, [], [f8])).
fof(f5077, plain, ((e21 = e22) | (~ spl102_90 | ~ spl102_102 | ~ spl102_233)), inference(forward_demodulation, [], [f5076, f1437])).
fof(f5076, plain, ((e22 = op2(e22, e21)) | (~ spl102_102 | ~ spl102_233)), inference(forward_demodulation, [], [f2350, f1488])).
fof(f2350, plain, ((e22 = op2(e22, op2(e21, e22))) | ~ spl102_233), inference(avatar_component_clause, [], [f2348])).
fof(f2348, plain, (spl102_233 <=> (e22 = op2(e22, op2(e21, e22)))), introduced(avatar_definition, [new_symbols(naming, [spl102_233])])).
fof(f5072, plain, (spl102_106 | ~ spl102_292 | ~ spl102_293), inference(avatar_contradiction_clause, [], [f5071])).
fof(f5071, plain, ($false | (spl102_106 | ~ spl102_292 | ~ spl102_293)), inference(subsumption_resolution, [], [f5069, f1504])).
fof(f5069, plain, ((e21 = op2(e21, e21)) | (~ spl102_292 | ~ spl102_293)), inference(backward_demodulation, [], [f4777, f2775])).
fof(f4777, plain, ((e21 = op2(h2(e13), e21)) | ~ spl102_293), inference(forward_demodulation, [], [f2605, f2780])).
fof(f2780, plain, ((e21 = h2(e12)) | ~ spl102_293), inference(avatar_component_clause, [], [f2779])).
fof(f5038, plain, (spl102_292 | ~ spl102_102 | ~ spl102_290), inference(avatar_split_clause, [], [f5000, f2764, f1486, f2774])).
fof(f5000, plain, ((e21 = h2(e13)) | (~ spl102_102 | ~ spl102_290)), inference(backward_demodulation, [], [f4924, f1488])).
fof(f4924, plain, ((op2(e21, e22) = h2(e13)) | ~ spl102_290), inference(backward_demodulation, [], [f2612, f2765])).
fof(f4990, plain, (~ spl102_83 | ~ spl102_115), inference(avatar_split_clause, [], [f4989, f1541, f1405])).
fof(f1405, plain, (spl102_83 <=> (e22 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl102_83])])).
fof(f4989, plain, (~ (e22 = op2(e22, e23)) | ~ spl102_115), inference(backward_demodulation, [], [f393, f1543])).
fof(f393, plain, ~ (op2(e20, e23) = op2(e22, e23)), inference(cnf_transformation, [], [f6])).
fof(f4966, plain, (~ spl102_70 | ~ spl102_270), inference(avatar_split_clause, [], [f4959, f2663, f1350])).
fof(f1350, plain, (spl102_70 <=> (e21 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl102_70])])).
fof(f4959, plain, (~ (e21 = op2(e23, e22)) | ~ spl102_270), inference(backward_demodulation, [], [f2626, f2664])).
fof(f4928, plain, (~ spl102_103 | ~ spl102_290), inference(avatar_split_clause, [], [f4922, f2764, f1490])).
fof(f1490, plain, (spl102_103 <=> (e22 = op2(e21, e22))), introduced(avatar_definition, [new_symbols(naming, [spl102_103])])).
fof(f4922, plain, (~ (e22 = op2(e21, e22)) | ~ spl102_290), inference(backward_demodulation, [], [f2610, f2765])).
fof(f2610, plain, ~ (op2(e21, e22) = h2(e11)), inference(backward_demodulation, [], [f407, f909])).
fof(f407, plain, ~ (op2(e21, e21) = op2(e21, e22)), inference(cnf_transformation, [], [f6])).
fof(f4897, plain, (~ spl102_101 | ~ spl102_73 | ~ spl102_226), inference(avatar_split_clause, [], [f4896, f2315, f1363, f1482])).
fof(f1482, plain, (spl102_101 <=> (e20 = op2(e21, e22))), introduced(avatar_definition, [new_symbols(naming, [spl102_101])])).
fof(f4896, plain, (~ (e20 = op2(e21, e22)) | (~ spl102_73 | ~ spl102_226)), inference(forward_demodulation, [], [f4769, f1365])).
fof(f4769, plain, (~ (op2(e21, e22) = op2(e23, e21)) | ~ spl102_226), inference(forward_demodulation, [], [f385, f2317])).
fof(f2317, plain, ((op2(e21, e22) = op2(e22, e21)) | ~ spl102_226), inference(avatar_component_clause, [], [f2315])).
fof(f4877, plain, (~ spl102_1 | ~ spl102_9), inference(avatar_split_clause, [], [f4872, f1059, f1025])).
fof(f1025, plain, (spl102_1 <=> (e10 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl102_1])])).
fof(f4872, plain, (~ (e10 = op1(e13, e13)) | ~ spl102_9), inference(backward_demodulation, [], [f372, f1061])).
fof(f372, plain, ~ (op1(e13, e11) = op1(e13, e13)), inference(cnf_transformation, [], [f5])).
fof(f4862, plain, (~ spl102_10 | ~ spl102_26), inference(avatar_split_clause, [], [f4859, f1131, f1063])).
fof(f4859, plain, (~ (e11 = op1(e13, e11)) | ~ spl102_26), inference(backward_demodulation, [], [f337, f1133])).
fof(f4861, plain, (~ spl102_18 | ~ spl102_26), inference(avatar_split_clause, [], [f4858, f1131, f1097])).
fof(f1097, plain, (spl102_18 <=> (e11 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl102_18])])).
fof(f4858, plain, (~ (e11 = op1(e12, e13)) | ~ spl102_26), inference(backward_demodulation, [], [f366, f1133])).
fof(f4833, plain, (~ spl102_69 | ~ spl102_73), inference(avatar_split_clause, [], [f4832, f1363, f1346])).
fof(f1346, plain, (spl102_69 <=> (e20 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl102_69])])).
fof(f4832, plain, (~ (e20 = op2(e23, e22)) | ~ spl102_73), inference(backward_demodulation, [], [f419, f1365])).
fof(f419, plain, ~ (op2(e23, e21) = op2(e23, e22)), inference(cnf_transformation, [], [f6])).
fof(f4826, plain, (~ spl102_88 | spl102_345), inference(avatar_contradiction_clause, [], [f4825])).
fof(f4825, plain, ($false | (~ spl102_88 | spl102_345)), inference(subsumption_resolution, [], [f4824, f3087])).
fof(f3087, plain, (~ (e23 = h3(e11)) | spl102_345), inference(avatar_component_clause, [], [f3085])).
fof(f4824, plain, ((e23 = h3(e11)) | ~ spl102_88), inference(backward_demodulation, [], [f913, f1428])).
fof(f1428, plain, ((e23 = op2(e22, e22)) | ~ spl102_88), inference(avatar_component_clause, [], [f1426])).
fof(f1426, plain, (spl102_88 <=> (e23 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl102_88])])).
fof(f4823, plain, (~ spl102_90 | ~ spl102_91), inference(avatar_contradiction_clause, [], [f4822])).
fof(f4822, plain, ($false | (~ spl102_90 | ~ spl102_91)), inference(subsumption_resolution, [], [f4821, f431])).
fof(f4821, plain, ((e21 = e22) | (~ spl102_90 | ~ spl102_91)), inference(backward_demodulation, [], [f1441, f1437])).
fof(f1441, plain, ((e22 = op2(e22, e21)) | ~ spl102_91), inference(avatar_component_clause, [], [f1439])).
fof(f1439, plain, (spl102_91 <=> (e22 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl102_91])])).
fof(f4804, plain, (spl102_91 | ~ spl102_103 | ~ spl102_226), inference(avatar_split_clause, [], [f4797, f2315, f1490, f1439])).
fof(f4797, plain, ((e22 = op2(e22, e21)) | (~ spl102_103 | ~ spl102_226)), inference(backward_demodulation, [], [f2317, f1492])).
fof(f1492, plain, ((e22 = op2(e21, e22)) | ~ spl102_103), inference(avatar_component_clause, [], [f1490])).
fof(f4794, plain, (~ spl102_106 | spl102_294), inference(avatar_contradiction_clause, [], [f4793])).
fof(f4793, plain, ($false | (~ spl102_106 | spl102_294)), inference(subsumption_resolution, [], [f4791, f2786])).
fof(f2786, plain, (~ (e21 = h2(e11)) | spl102_294), inference(avatar_component_clause, [], [f2784])).
fof(f2784, plain, (spl102_294 <=> (e21 = h2(e11))), introduced(avatar_definition, [new_symbols(naming, [spl102_294])])).
fof(f4791, plain, ((e21 = h2(e11)) | ~ spl102_106), inference(backward_demodulation, [], [f909, f1505])).
fof(f1505, plain, ((e21 = op2(e21, e21)) | ~ spl102_106), inference(avatar_component_clause, [], [f1503])).
fof(f4787, plain, (~ spl102_96 | ~ spl102_112), inference(avatar_split_clause, [], [f4784, f1528, f1460])).
fof(f4784, plain, (~ (e23 = op2(e22, e20)) | ~ spl102_112), inference(backward_demodulation, [], [f377, f1530])).
fof(f4786, plain, (~ spl102_104 | ~ spl102_112), inference(avatar_split_clause, [], [f4783, f1528, f1494])).
fof(f1494, plain, (spl102_104 <=> (e23 = op2(e21, e22))), introduced(avatar_definition, [new_symbols(naming, [spl102_104])])).
fof(f4783, plain, (~ (e23 = op2(e21, e22)) | ~ spl102_112), inference(backward_demodulation, [], [f405, f1530])).
fof(f405, plain, ~ (op2(e21, e20) = op2(e21, e22)), inference(cnf_transformation, [], [f6])).
fof(f4779, plain, (~ spl102_74 | ~ spl102_98 | spl102_216), inference(avatar_split_clause, [], [f4778, f2267, f1469, f1367])).
fof(f1469, plain, (spl102_98 <=> (e21 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl102_98])])).
fof(f4778, plain, (~ (e21 = op2(e23, e21)) | (~ spl102_98 | spl102_216)), inference(forward_demodulation, [], [f2268, f1471])).
fof(f1471, plain, ((e21 = op2(e21, e23)) | ~ spl102_98), inference(avatar_component_clause, [], [f1469])).
fof(f2268, plain, (~ (op2(e21, e23) = op2(e23, e21)) | spl102_216), inference(avatar_component_clause, [], [f2267])).
fof(f4776, plain, (~ spl102_103 | ~ spl102_119), inference(avatar_split_clause, [], [f4775, f1558, f1490])).
fof(f4775, plain, (~ (e22 = op2(e21, e22)) | ~ spl102_119), inference(forward_demodulation, [], [f386, f1560])).
fof(f386, plain, ~ (op2(e20, e22) = op2(e21, e22)), inference(cnf_transformation, [], [f6])).
fof(f4752, plain, (~ spl102_1 | ~ spl102_4), inference(avatar_contradiction_clause, [], [f4751])).
fof(f4751, plain, ($false | (~ spl102_1 | ~ spl102_4)), inference(subsumption_resolution, [], [f4750, f424])).
fof(f424, plain, ~ (e10 = e13), inference(cnf_transformation, [], [f7])).
fof(f4750, plain, ((e10 = e13) | (~ spl102_1 | ~ spl102_4)), inference(forward_demodulation, [], [f1039, f1027])).
fof(f1027, plain, ((e10 = op1(e13, e13)) | ~ spl102_1), inference(avatar_component_clause, [], [f1025])).
fof(f4728, plain, (~ spl102_89 | spl102_101 | ~ spl102_226), inference(avatar_contradiction_clause, [], [f4727])).
fof(f4727, plain, ($false | (~ spl102_89 | spl102_101 | ~ spl102_226)), inference(subsumption_resolution, [], [f4726, f1483])).
fof(f1483, plain, (~ (e20 = op2(e21, e22)) | spl102_101), inference(avatar_component_clause, [], [f1482])).
fof(f4726, plain, ((e20 = op2(e21, e22)) | (~ spl102_89 | ~ spl102_226)), inference(forward_demodulation, [], [f2317, f1433])).
fof(f1433, plain, ((e20 = op2(e22, e21)) | ~ spl102_89), inference(avatar_component_clause, [], [f1431])).
fof(f1431, plain, (spl102_89 <=> (e20 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl102_89])])).
fof(f4705, plain, (~ spl102_34 | ~ spl102_18), inference(avatar_split_clause, [], [f4704, f1097, f1165])).
fof(f4704, plain, (~ (e11 = op1(e11, e13)) | ~ spl102_18), inference(forward_demodulation, [], [f347, f1099])).
fof(f1099, plain, ((e11 = op1(e12, e13)) | ~ spl102_18), inference(avatar_component_clause, [], [f1097])).
fof(f347, plain, ~ (op1(e11, e13) = op1(e12, e13)), inference(cnf_transformation, [], [f5])).
fof(f4702, plain, (~ spl102_34 | ~ spl102_10 | spl102_149), inference(avatar_split_clause, [], [f4701, f1791, f1063, f1165])).
fof(f4701, plain, (~ (e11 = op1(e11, e13)) | (~ spl102_10 | spl102_149)), inference(forward_demodulation, [], [f1792, f1065])).
fof(f4687, plain, (~ spl102_25 | spl102_37 | ~ spl102_159), inference(avatar_contradiction_clause, [], [f4686])).
fof(f4686, plain, ($false | (~ spl102_25 | spl102_37 | ~ spl102_159)), inference(subsumption_resolution, [], [f4685, f1179])).
fof(f1179, plain, (~ (e10 = op1(e11, e12)) | spl102_37), inference(avatar_component_clause, [], [f1178])).
fof(f1178, plain, (spl102_37 <=> (e10 = op1(e11, e12))), introduced(avatar_definition, [new_symbols(naming, [spl102_37])])).
fof(f4685, plain, ((e10 = op1(e11, e12)) | (~ spl102_25 | ~ spl102_159)), inference(forward_demodulation, [], [f1841, f1129])).
fof(f1841, plain, ((op1(e11, e12) = op1(e12, e11)) | ~ spl102_159), inference(avatar_component_clause, [], [f1839])).
fof(f1839, plain, (spl102_159 <=> (op1(e11, e12) = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl102_159])])).
fof(f4664, plain, (~ spl102_2 | ~ spl102_18), inference(avatar_split_clause, [], [f4660, f1097, f1029])).
fof(f4660, plain, (~ (e11 = op1(e13, e13)) | ~ spl102_18), inference(backward_demodulation, [], [f349, f1099])).
fof(f349, plain, ~ (op1(e12, e13) = op1(e13, e13)), inference(cnf_transformation, [], [f5])).
fof(f4620, plain, (~ spl102_109 | spl102_125 | ~ spl102_237), inference(avatar_contradiction_clause, [], [f4619])).
fof(f4619, plain, ($false | (~ spl102_109 | spl102_125 | ~ spl102_237)), inference(subsumption_resolution, [], [f4618, f1585])).
fof(f1585, plain, (~ (e20 = op2(e20, e20)) | spl102_125), inference(avatar_component_clause, [], [f1584])).
fof(f1584, plain, (spl102_125 <=> (e20 = op2(e20, e20))), introduced(avatar_definition, [new_symbols(naming, [spl102_125])])).
fof(f4618, plain, ((e20 = op2(e20, e20)) | (~ spl102_109 | ~ spl102_237)), inference(forward_demodulation, [], [f2370, f1518])).
fof(f2370, plain, ((e20 = op2(e20, op2(e21, e20))) | ~ spl102_237), inference(avatar_component_clause, [], [f2368])).
fof(f2368, plain, (spl102_237 <=> (e20 = op2(e20, op2(e21, e20)))), introduced(avatar_definition, [new_symbols(naming, [spl102_237])])).
fof(f4617, plain, (~ spl102_109 | spl102_121 | ~ spl102_238), inference(avatar_contradiction_clause, [], [f4616])).
fof(f4616, plain, ($false | (~ spl102_109 | spl102_121 | ~ spl102_238)), inference(subsumption_resolution, [], [f4615, f1568])).
fof(f1568, plain, (~ (e20 = op2(e20, e21)) | spl102_121), inference(avatar_component_clause, [], [f1567])).
fof(f1567, plain, (spl102_121 <=> (e20 = op2(e20, e21))), introduced(avatar_definition, [new_symbols(naming, [spl102_121])])).
fof(f4615, plain, ((e20 = op2(e20, e21)) | (~ spl102_109 | ~ spl102_238)), inference(forward_demodulation, [], [f2375, f1518])).
fof(f4590, plain, (~ spl102_26 | ~ spl102_38 | ~ spl102_166), inference(avatar_contradiction_clause, [], [f4589])).
fof(f4589, plain, ($false | (~ spl102_26 | ~ spl102_38 | ~ spl102_166)), inference(subsumption_resolution, [], [f4588, f425])).
fof(f425, plain, ~ (e11 = e12), inference(cnf_transformation, [], [f7])).
fof(f4588, plain, ((e11 = e12) | (~ spl102_26 | ~ spl102_38 | ~ spl102_166)), inference(forward_demodulation, [], [f4587, f1133])).
fof(f4587, plain, ((e12 = op1(e12, e11)) | (~ spl102_38 | ~ spl102_166)), inference(forward_demodulation, [], [f1874, f1184])).
fof(f1874, plain, ((e12 = op1(e12, op1(e11, e12))) | ~ spl102_166), inference(avatar_component_clause, [], [f1872])).
fof(f1872, plain, (spl102_166 <=> (e12 = op1(e12, op1(e11, e12)))), introduced(avatar_definition, [new_symbols(naming, [spl102_166])])).
fof(f4560, plain, (~ spl102_79 | ~ spl102_113 | ~ spl102_219), inference(avatar_contradiction_clause, [], [f4559])).
fof(f4559, plain, ($false | (~ spl102_79 | ~ spl102_113 | ~ spl102_219)), inference(subsumption_resolution, [], [f4558, f429])).
fof(f429, plain, ~ (e20 = e22), inference(cnf_transformation, [], [f8])).
fof(f4558, plain, ((e20 = e22) | (~ spl102_79 | ~ spl102_113 | ~ spl102_219)), inference(forward_demodulation, [], [f4557, f1535])).
fof(f1535, plain, ((e20 = op2(e20, e23)) | ~ spl102_113), inference(avatar_component_clause, [], [f1533])).
fof(f4557, plain, ((e22 = op2(e20, e23)) | (~ spl102_79 | ~ spl102_219)), inference(forward_demodulation, [], [f2283, f1390])).
fof(f2283, plain, ((op2(e20, e23) = op2(e23, e20)) | ~ spl102_219), inference(avatar_component_clause, [], [f2281])).
fof(f2281, plain, (spl102_219 <=> (op2(e20, e23) = op2(e23, e20))), introduced(avatar_definition, [new_symbols(naming, [spl102_219])])).
fof(f4548, plain, (spl102_326 | ~ spl102_68), inference(avatar_split_clause, [], [f4495, f1341, f2964])).
fof(f1341, plain, (spl102_68 <=> (e23 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl102_68])])).
fof(f4495, plain, ((e23 = h4(e11)) | ~ spl102_68), inference(forward_demodulation, [], [f917, f1343])).
fof(f1343, plain, ((e23 = op2(e23, e23)) | ~ spl102_68), inference(avatar_component_clause, [], [f1341])).
fof(f4528, plain, (~ spl102_26 | spl102_42 | ~ spl102_158), inference(avatar_contradiction_clause, [], [f4527])).
fof(f4527, plain, ($false | (~ spl102_26 | spl102_42 | ~ spl102_158)), inference(subsumption_resolution, [], [f4526, f1200])).
fof(f4526, plain, ((e11 = op1(e11, e11)) | (~ spl102_26 | ~ spl102_158)), inference(forward_demodulation, [], [f1836, f1133])).
fof(f1836, plain, ((e11 = op1(e11, op1(e12, e11))) | ~ spl102_158), inference(avatar_component_clause, [], [f1834])).
fof(f1834, plain, (spl102_158 <=> (e11 = op1(e11, op1(e12, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl102_158])])).
fof(f4518, plain, (~ spl102_74 | ~ spl102_98 | ~ spl102_231), inference(avatar_contradiction_clause, [], [f4517])).
fof(f4517, plain, ($false | (~ spl102_74 | ~ spl102_98 | ~ spl102_231)), inference(subsumption_resolution, [], [f4516, f432])).
fof(f4516, plain, ((e21 = e23) | (~ spl102_74 | ~ spl102_98 | ~ spl102_231)), inference(forward_demodulation, [], [f4515, f1369])).
fof(f1369, plain, ((e21 = op2(e23, e21)) | ~ spl102_74), inference(avatar_component_clause, [], [f1367])).
fof(f4515, plain, ((e23 = op2(e23, e21)) | (~ spl102_98 | ~ spl102_231)), inference(forward_demodulation, [], [f2340, f1471])).
fof(f2340, plain, ((e23 = op2(e23, op2(e21, e23))) | ~ spl102_231), inference(avatar_component_clause, [], [f2338])).
fof(f2338, plain, (spl102_231 <=> (e23 = op2(e23, op2(e21, e23)))), introduced(avatar_definition, [new_symbols(naming, [spl102_231])])).
fof(f4469, plain, (~ spl102_65 | ~ spl102_68), inference(avatar_contradiction_clause, [], [f4468])).
fof(f4468, plain, ($false | (~ spl102_65 | ~ spl102_68)), inference(subsumption_resolution, [], [f4467, f430])).
fof(f4467, plain, ((e20 = e23) | (~ spl102_65 | ~ spl102_68)), inference(forward_demodulation, [], [f1343, f1331])).
fof(f1331, plain, ((e20 = op2(e23, e23)) | ~ spl102_65), inference(avatar_component_clause, [], [f1329])).
fof(f1329, plain, (spl102_65 <=> (e20 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl102_65])])).
fof(f4449, plain, (~ spl102_69 | ~ spl102_117), inference(avatar_split_clause, [], [f4446, f1550, f1346])).
fof(f4446, plain, (~ (e20 = op2(e23, e22)) | ~ spl102_117), inference(backward_demodulation, [], [f388, f1552])).
fof(f388, plain, ~ (op2(e20, e22) = op2(e23, e22)), inference(cnf_transformation, [], [f6])).
fof(f4406, plain, (spl102_71 | ~ spl102_83 | ~ spl102_213), inference(avatar_split_clause, [], [f4405, f2253, f1405, f1354])).
fof(f1354, plain, (spl102_71 <=> (e22 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl102_71])])).
fof(f2253, plain, (spl102_213 <=> (op2(e22, e23) = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl102_213])])).
fof(f4405, plain, ((e22 = op2(e23, e22)) | (~ spl102_83 | ~ spl102_213)), inference(forward_demodulation, [], [f2255, f1407])).
fof(f1407, plain, ((e22 = op2(e22, e23)) | ~ spl102_83), inference(avatar_component_clause, [], [f1405])).
fof(f2255, plain, ((op2(e22, e23) = op2(e23, e22)) | ~ spl102_213), inference(avatar_component_clause, [], [f2253])).
fof(f4404, plain, (~ spl102_71 | ~ spl102_300), inference(avatar_split_clause, [], [f4096, f2816, f1354])).
fof(f4096, plain, (~ (e22 = op2(e23, e22)) | ~ spl102_300), inference(forward_demodulation, [], [f2596, f2817])).
fof(f2596, plain, ~ (op2(e23, e22) = h1(e12)), inference(backward_demodulation, [], [f417, f2591])).
fof(f417, plain, ~ (op2(e23, e20) = op2(e23, e22)), inference(cnf_transformation, [], [f6])).
fof(f4394, plain, (~ spl102_35 | ~ spl102_43), inference(avatar_split_clause, [], [f4393, f1203, f1169])).
fof(f1169, plain, (spl102_35 <=> (e12 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl102_35])])).
fof(f4393, plain, (~ (e12 = op1(e11, e13)) | ~ spl102_43), inference(forward_demodulation, [], [f360, f1205])).
fof(f360, plain, ~ (op1(e11, e11) = op1(e11, e13)), inference(cnf_transformation, [], [f5])).
fof(f4385, plain, (spl102_20 | ~ spl102_8 | ~ spl102_146), inference(avatar_split_clause, [], [f3851, f1777, f1054, f1105])).
fof(f3851, plain, ((e13 = op1(e12, e13)) | (~ spl102_8 | ~ spl102_146)), inference(forward_demodulation, [], [f1779, f1056])).
fof(f4379, plain, (~ spl102_37 | ~ spl102_38), inference(avatar_contradiction_clause, [], [f4378])).
fof(f4378, plain, ($false | (~ spl102_37 | ~ spl102_38)), inference(subsumption_resolution, [], [f4377, f422])).
fof(f422, plain, ~ (e10 = e11), inference(cnf_transformation, [], [f7])).
fof(f4377, plain, ((e10 = e11) | (~ spl102_37 | ~ spl102_38)), inference(backward_demodulation, [], [f1184, f1180])).
fof(f1180, plain, ((e10 = op1(e11, e12)) | ~ spl102_37), inference(avatar_component_clause, [], [f1178])).
fof(f4376, plain, (~ spl102_62 | ~ spl102_64), inference(avatar_contradiction_clause, [], [f4375])).
fof(f4375, plain, ($false | (~ spl102_62 | ~ spl102_64)), inference(subsumption_resolution, [], [f4374, f426])).
fof(f4374, plain, ((e11 = e13) | (~ spl102_62 | ~ spl102_64)), inference(forward_demodulation, [], [f1294, f1286])).
fof(f1294, plain, ((op1(e10, e10) = e13) | ~ spl102_64), inference(avatar_component_clause, [], [f1292])).
fof(f1292, plain, (spl102_64 <=> (op1(e10, e10) = e13)), introduced(avatar_definition, [new_symbols(naming, [spl102_64])])).
fof(f4366, plain, (~ spl102_73 | ~ spl102_74), inference(avatar_contradiction_clause, [], [f4365])).
fof(f4365, plain, ($false | (~ spl102_73 | ~ spl102_74)), inference(subsumption_resolution, [], [f4364, f428])).
fof(f4364, plain, ((e20 = e21) | (~ spl102_73 | ~ spl102_74)), inference(backward_demodulation, [], [f1369, f1365])).
fof(f4354, plain, (~ spl102_86 | ~ spl102_88), inference(avatar_contradiction_clause, [], [f4353])).
fof(f4353, plain, ($false | (~ spl102_86 | ~ spl102_88)), inference(subsumption_resolution, [], [f4352, f432])).
fof(f4352, plain, ((e21 = e23) | (~ spl102_86 | ~ spl102_88)), inference(backward_demodulation, [], [f1428, f1420])).
fof(f1420, plain, ((e21 = op2(e22, e22)) | ~ spl102_86), inference(avatar_component_clause, [], [f1418])).
fof(f1418, plain, (spl102_86 <=> (e21 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl102_86])])).
fof(f4323, plain, (~ spl102_15 | ~ spl102_49 | ~ spl102_152), inference(avatar_contradiction_clause, [], [f4322])).
fof(f4322, plain, ($false | (~ spl102_15 | ~ spl102_49 | ~ spl102_152)), inference(subsumption_resolution, [], [f4321, f423])).
fof(f423, plain, ~ (e10 = e12), inference(cnf_transformation, [], [f7])).
fof(f4321, plain, ((e10 = e12) | (~ spl102_15 | ~ spl102_49 | ~ spl102_152)), inference(forward_demodulation, [], [f4320, f1231])).
fof(f4320, plain, ((e12 = op1(e10, e13)) | (~ spl102_15 | ~ spl102_152)), inference(forward_demodulation, [], [f1807, f1086])).
fof(f4314, plain, (~ spl102_73 | ~ spl102_274), inference(avatar_split_clause, [], [f4307, f2683, f1363])).
fof(f2683, plain, (spl102_274 <=> (e20 = h4(e11))), introduced(avatar_definition, [new_symbols(naming, [spl102_274])])).
fof(f4307, plain, (~ (e20 = op2(e23, e21)) | ~ spl102_274), inference(backward_demodulation, [], [f2625, f2684])).
fof(f2684, plain, ((e20 = h4(e11)) | ~ spl102_274), inference(avatar_component_clause, [], [f2683])).
fof(f2625, plain, ~ (op2(e23, e21) = h4(e11)), inference(backward_demodulation, [], [f420, f917])).
fof(f420, plain, ~ (op2(e23, e21) = op2(e23, e23)), inference(cnf_transformation, [], [f6])).
fof(f4312, plain, (~ spl102_113 | ~ spl102_274), inference(avatar_split_clause, [], [f4305, f2683, f1533])).
fof(f4305, plain, (~ (e20 = op2(e20, e23)) | ~ spl102_274), inference(backward_demodulation, [], [f2622, f2684])).
fof(f2622, plain, ~ (op2(e20, e23) = h4(e11)), inference(backward_demodulation, [], [f394, f917])).
fof(f394, plain, ~ (op2(e20, e23) = op2(e23, e23)), inference(cnf_transformation, [], [f6])).
fof(f4298, plain, (~ spl102_90 | ~ spl102_282), inference(avatar_split_clause, [], [f4291, f2724, f1435])).
fof(f2724, plain, (spl102_282 <=> (e21 = h3(e11))), introduced(avatar_definition, [new_symbols(naming, [spl102_282])])).
fof(f4291, plain, (~ (e21 = op2(e22, e21)) | ~ spl102_282), inference(backward_demodulation, [], [f2618, f2725])).
fof(f2725, plain, ((e21 = h3(e11)) | ~ spl102_282), inference(avatar_component_clause, [], [f2724])).
fof(f2618, plain, ~ (op2(e22, e21) = h3(e11)), inference(backward_demodulation, [], [f413, f913])).
fof(f413, plain, ~ (op2(e22, e21) = op2(e22, e22)), inference(cnf_transformation, [], [f6])).
fof(f4268, plain, (~ spl102_92 | ~ spl102_124), inference(avatar_split_clause, [], [f4267, f1579, f1443])).
fof(f1443, plain, (spl102_92 <=> (e23 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl102_92])])).
fof(f4267, plain, (~ (e23 = op2(e22, e21)) | ~ spl102_124), inference(backward_demodulation, [], [f381, f1581])).
fof(f381, plain, ~ (op2(e20, e21) = op2(e22, e21)), inference(cnf_transformation, [], [f6])).
fof(f4266, plain, (~ spl102_116 | ~ spl102_124), inference(avatar_split_clause, [], [f3857, f1579, f1545])).
fof(f1545, plain, (spl102_116 <=> (e23 = op2(e20, e23))), introduced(avatar_definition, [new_symbols(naming, [spl102_116])])).
fof(f3857, plain, (~ (e23 = op2(e20, e23)) | ~ spl102_124), inference(backward_demodulation, [], [f402, f1581])).
fof(f402, plain, ~ (op2(e20, e21) = op2(e20, e23)), inference(cnf_transformation, [], [f6])).
fof(f4258, plain, (~ spl102_90 | ~ spl102_98 | ~ spl102_216), inference(avatar_split_clause, [], [f4257, f2267, f1469, f1435])).
fof(f4257, plain, (~ (e21 = op2(e22, e21)) | (~ spl102_98 | ~ spl102_216)), inference(forward_demodulation, [], [f3794, f1471])).
fof(f3794, plain, (~ (op2(e21, e23) = op2(e22, e21)) | ~ spl102_216), inference(forward_demodulation, [], [f385, f2269])).
fof(f4250, plain, (~ spl102_101 | ~ spl102_69), inference(avatar_split_clause, [], [f4249, f1346, f1482])).
fof(f4249, plain, (~ (e20 = op2(e21, e22)) | ~ spl102_69), inference(forward_demodulation, [], [f390, f1348])).
fof(f1348, plain, ((e20 = op2(e23, e22)) | ~ spl102_69), inference(avatar_component_clause, [], [f1346])).
fof(f390, plain, ~ (op2(e21, e22) = op2(e23, e22)), inference(cnf_transformation, [], [f6])).
fof(f4244, plain, (~ spl102_42 | ~ spl102_38), inference(avatar_split_clause, [], [f4074, f1182, f1199])).
fof(f4074, plain, (~ (e11 = op1(e11, e11)) | ~ spl102_38), inference(forward_demodulation, [], [f359, f1184])).
fof(f359, plain, ~ (op1(e11, e11) = op1(e11, e12)), inference(cnf_transformation, [], [f5])).
fof(f4240, plain, (~ spl102_41 | ~ spl102_9), inference(avatar_split_clause, [], [f4084, f1059, f1195])).
fof(f1195, plain, (spl102_41 <=> (e10 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl102_41])])).
fof(f4084, plain, (~ (e10 = op1(e11, e11)) | ~ spl102_9), inference(forward_demodulation, [], [f336, f1061])).
fof(f336, plain, ~ (op1(e11, e11) = op1(e13, e11)), inference(cnf_transformation, [], [f5])).
fof(f4239, plain, (~ spl102_42 | ~ spl102_26), inference(avatar_split_clause, [], [f4238, f1131, f1199])).
fof(f4238, plain, (~ (e11 = op1(e11, e11)) | ~ spl102_26), inference(forward_demodulation, [], [f335, f1133])).
fof(f335, plain, ~ (op1(e11, e11) = op1(e12, e11)), inference(cnf_transformation, [], [f5])).
fof(f4223, plain, (~ spl102_21 | ~ spl102_22), inference(avatar_contradiction_clause, [], [f4222])).
fof(f4222, plain, ($false | (~ spl102_21 | ~ spl102_22)), inference(subsumption_resolution, [], [f4221, f422])).
fof(f4221, plain, ((e10 = e11) | (~ spl102_21 | ~ spl102_22)), inference(backward_demodulation, [], [f1116, f1112])).
fof(f1112, plain, ((e10 = op1(e12, e12)) | ~ spl102_21), inference(avatar_component_clause, [], [f1110])).
fof(f1116, plain, ((e11 = op1(e12, e12)) | ~ spl102_22), inference(avatar_component_clause, [], [f1114])).
fof(f1114, plain, (spl102_22 <=> (e11 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl102_22])])).
fof(f4220, plain, (~ spl102_22 | ~ spl102_23), inference(avatar_contradiction_clause, [], [f4219])).
fof(f4219, plain, ($false | (~ spl102_22 | ~ spl102_23)), inference(subsumption_resolution, [], [f4218, f425])).
fof(f4218, plain, ((e11 = e12) | (~ spl102_22 | ~ spl102_23)), inference(forward_demodulation, [], [f1120, f1116])).
fof(f4214, plain, (~ spl102_27 | ~ spl102_35 | ~ spl102_149), inference(avatar_split_clause, [], [f4210, f1791, f1169, f1135])).
fof(f1135, plain, (spl102_27 <=> (e12 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl102_27])])).
fof(f4210, plain, (~ (e12 = op1(e12, e11)) | (~ spl102_35 | ~ spl102_149)), inference(backward_demodulation, [], [f3921, f1171])).
fof(f1171, plain, ((e12 = op1(e11, e13)) | ~ spl102_35), inference(avatar_component_clause, [], [f1169])).
fof(f3921, plain, (~ (op1(e11, e13) = op1(e12, e11)) | ~ spl102_149), inference(backward_demodulation, [], [f337, f1793])).
fof(f4206, plain, (~ spl102_41 | ~ spl102_43), inference(avatar_contradiction_clause, [], [f4205])).
fof(f4205, plain, ($false | (~ spl102_41 | ~ spl102_43)), inference(subsumption_resolution, [], [f4204, f423])).
fof(f4204, plain, ((e10 = e12) | (~ spl102_41 | ~ spl102_43)), inference(forward_demodulation, [], [f1205, f1197])).
fof(f1197, plain, ((e10 = op1(e11, e11)) | ~ spl102_41), inference(avatar_component_clause, [], [f1195])).
fof(f4186, plain, (~ spl102_61 | ~ spl102_62), inference(avatar_contradiction_clause, [], [f4185])).
fof(f4185, plain, ($false | (~ spl102_61 | ~ spl102_62)), inference(subsumption_resolution, [], [f4184, f422])).
fof(f4184, plain, ((e10 = e11) | (~ spl102_61 | ~ spl102_62)), inference(backward_demodulation, [], [f1286, f1282])).
fof(f1282, plain, ((e10 = op1(e10, e10)) | ~ spl102_61), inference(avatar_component_clause, [], [f1280])).
fof(f4172, plain, (~ spl102_84 | ~ spl102_96), inference(avatar_split_clause, [], [f4171, f1460, f1409])).
fof(f1409, plain, (spl102_84 <=> (e23 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl102_84])])).
fof(f4171, plain, (~ (e23 = op2(e22, e23)) | ~ spl102_96), inference(backward_demodulation, [], [f412, f1462])).
fof(f412, plain, ~ (op2(e22, e20) = op2(e22, e23)), inference(cnf_transformation, [], [f6])).
fof(f4156, plain, (~ spl102_121 | ~ spl102_124), inference(avatar_contradiction_clause, [], [f4155])).
fof(f4155, plain, ($false | (~ spl102_121 | ~ spl102_124)), inference(subsumption_resolution, [], [f4154, f430])).
fof(f4154, plain, ((e20 = e23) | (~ spl102_121 | ~ spl102_124)), inference(backward_demodulation, [], [f1581, f1569])).
fof(f1569, plain, ((e20 = op2(e20, e21)) | ~ spl102_121), inference(avatar_component_clause, [], [f1567])).
fof(f4134, plain, (~ spl102_111 | ~ spl102_300), inference(avatar_split_clause, [], [f4133, f2816, f1524])).
fof(f1524, plain, (spl102_111 <=> (e22 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl102_111])])).
fof(f4133, plain, (~ (e22 = op2(e21, e20)) | ~ spl102_300), inference(forward_demodulation, [], [f2593, f2817])).
fof(f2593, plain, ~ (op2(e21, e20) = h1(e12)), inference(backward_demodulation, [], [f378, f2591])).
fof(f378, plain, ~ (op2(e21, e20) = op2(e23, e20)), inference(cnf_transformation, [], [f6])).
fof(f4123, plain, (~ spl102_102 | ~ spl102_98), inference(avatar_split_clause, [], [f4122, f1469, f1486])).
fof(f4122, plain, (~ (e21 = op2(e21, e22)) | ~ spl102_98), inference(forward_demodulation, [], [f409, f1471])).
fof(f409, plain, ~ (op2(e21, e22) = op2(e21, e23)), inference(cnf_transformation, [], [f6])).
fof(f4119, plain, (~ spl102_294 | ~ spl102_98), inference(avatar_split_clause, [], [f4118, f1469, f2784])).
fof(f4118, plain, (~ (e21 = h2(e11)) | ~ spl102_98), inference(forward_demodulation, [], [f2611, f1471])).
fof(f4105, plain, (~ spl102_74 | ~ spl102_270), inference(avatar_split_clause, [], [f3952, f2663, f1367])).
fof(f3952, plain, (~ (e21 = op2(e23, e21)) | ~ spl102_270), inference(forward_demodulation, [], [f2625, f2664])).
fof(f4104, plain, (~ spl102_76 | ~ spl102_124), inference(avatar_split_clause, [], [f4103, f1579, f1375])).
fof(f4103, plain, (~ (e23 = op2(e23, e21)) | ~ spl102_124), inference(forward_demodulation, [], [f382, f1581])).
fof(f382, plain, ~ (op2(e20, e21) = op2(e23, e21)), inference(cnf_transformation, [], [f6])).
fof(f4093, plain, (~ spl102_102 | ~ spl102_70), inference(avatar_split_clause, [], [f4092, f1350, f1486])).
fof(f4092, plain, (~ (e21 = op2(e21, e22)) | ~ spl102_70), inference(forward_demodulation, [], [f390, f1352])).
fof(f1352, plain, ((e21 = op2(e23, e22)) | ~ spl102_70), inference(avatar_component_clause, [], [f1350])).
fof(f4081, plain, (~ spl102_22 | ~ spl102_38), inference(avatar_contradiction_clause, [], [f4080])).
fof(f4080, plain, ($false | (~ spl102_22 | ~ spl102_38)), inference(subsumption_resolution, [], [f4079, f1184])).
fof(f4079, plain, (~ (e11 = op1(e11, e12)) | ~ spl102_22), inference(forward_demodulation, [], [f341, f1116])).
fof(f341, plain, ~ (op1(e11, e12) = op1(e12, e12)), inference(cnf_transformation, [], [f5])).
fof(f4059, plain, (~ spl102_9 | ~ spl102_10), inference(avatar_contradiction_clause, [], [f4058])).
fof(f4058, plain, ($false | (~ spl102_9 | ~ spl102_10)), inference(subsumption_resolution, [], [f4057, f422])).
fof(f4057, plain, ((e10 = e11) | (~ spl102_9 | ~ spl102_10)), inference(backward_demodulation, [], [f1065, f1061])).
fof(f4054, plain, (~ spl102_18 | ~ spl102_20), inference(avatar_contradiction_clause, [], [f4053])).
fof(f4053, plain, ($false | (~ spl102_18 | ~ spl102_20)), inference(subsumption_resolution, [], [f4052, f426])).
fof(f4052, plain, ((e11 = e13) | (~ spl102_18 | ~ spl102_20)), inference(backward_demodulation, [], [f1107, f1099])).
fof(f1107, plain, ((e13 = op1(e12, e13)) | ~ spl102_20), inference(avatar_component_clause, [], [f1105])).
fof(f4045, plain, (spl102_12 | ~ spl102_36 | ~ spl102_149), inference(avatar_contradiction_clause, [], [f4044])).
fof(f4044, plain, ($false | (spl102_12 | ~ spl102_36 | ~ spl102_149)), inference(subsumption_resolution, [], [f4040, f1072])).
fof(f4040, plain, ((e13 = op1(e13, e11)) | (~ spl102_36 | ~ spl102_149)), inference(backward_demodulation, [], [f1793, f1175])).
fof(f4038, plain, (~ spl102_38 | ~ spl102_39), inference(avatar_contradiction_clause, [], [f4037])).
fof(f4037, plain, ($false | (~ spl102_38 | ~ spl102_39)), inference(subsumption_resolution, [], [f4036, f425])).
fof(f4036, plain, ((e11 = e12) | (~ spl102_38 | ~ spl102_39)), inference(backward_demodulation, [], [f1188, f1184])).
fof(f1188, plain, ((e12 = op1(e11, e12)) | ~ spl102_39), inference(avatar_component_clause, [], [f1186])).
fof(f4020, plain, (~ spl102_70 | ~ spl102_72), inference(avatar_contradiction_clause, [], [f4019])).
fof(f4019, plain, ($false | (~ spl102_70 | ~ spl102_72)), inference(subsumption_resolution, [], [f4018, f432])).
fof(f4018, plain, ((e21 = e23) | (~ spl102_70 | ~ spl102_72)), inference(backward_demodulation, [], [f1360, f1352])).
fof(f4016, plain, (~ spl102_77 | ~ spl102_79), inference(avatar_contradiction_clause, [], [f4015])).
fof(f4015, plain, ($false | (~ spl102_77 | ~ spl102_79)), inference(subsumption_resolution, [], [f4014, f429])).
fof(f4014, plain, ((e20 = e22) | (~ spl102_77 | ~ spl102_79)), inference(backward_demodulation, [], [f1390, f1382])).
fof(f1382, plain, ((e20 = op2(e23, e20)) | ~ spl102_77), inference(avatar_component_clause, [], [f1380])).
fof(f1380, plain, (spl102_77 <=> (e20 = op2(e23, e20))), introduced(avatar_definition, [new_symbols(naming, [spl102_77])])).
fof(f3993, plain, (~ spl102_103 | ~ spl102_278), inference(avatar_split_clause, [], [f3989, f2703, f1490])).
fof(f3989, plain, (~ (e22 = op2(e21, e22)) | ~ spl102_278), inference(backward_demodulation, [], [f2615, f2704])).
fof(f2615, plain, ~ (op2(e21, e22) = h3(e11)), inference(backward_demodulation, [], [f389, f913])).
fof(f389, plain, ~ (op2(e21, e22) = op2(e22, e22)), inference(cnf_transformation, [], [f6])).
fof(f3970, plain, (~ spl102_364 | ~ spl102_124), inference(avatar_split_clause, [], [f3969, f1579, f3195])).
fof(f3195, plain, (spl102_364 <=> (e23 = h2(e11))), introduced(avatar_definition, [new_symbols(naming, [spl102_364])])).
fof(f3969, plain, (~ (e23 = h2(e11)) | ~ spl102_124), inference(forward_demodulation, [], [f2606, f1581])).
fof(f2606, plain, ~ (op2(e20, e21) = h2(e11)), inference(backward_demodulation, [], [f380, f909])).
fof(f380, plain, ~ (op2(e20, e21) = op2(e21, e21)), inference(cnf_transformation, [], [f6])).
fof(f3966, plain, (~ spl102_298 | ~ spl102_97), inference(avatar_split_clause, [], [f3965, f1465, f2805])).
fof(f1465, plain, (spl102_97 <=> (e20 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl102_97])])).
fof(f3965, plain, (~ (e20 = h2(e11)) | ~ spl102_97), inference(forward_demodulation, [], [f2611, f1467])).
fof(f1467, plain, ((e20 = op2(e21, e23)) | ~ spl102_97), inference(avatar_component_clause, [], [f1465])).
fof(f3964, plain, (~ spl102_113 | ~ spl102_97), inference(avatar_split_clause, [], [f3963, f1465, f1533])).
fof(f3963, plain, (~ (e20 = op2(e20, e23)) | ~ spl102_97), inference(forward_demodulation, [], [f392, f1467])).
fof(f392, plain, ~ (op2(e20, e23) = op2(e21, e23)), inference(cnf_transformation, [], [f6])).
fof(f3951, plain, (spl102_272 | ~ spl102_73 | ~ spl102_270), inference(avatar_split_clause, [], [f3950, f2663, f1363, f2673])).
fof(f3950, plain, ((e20 = h4(e13)) | (~ spl102_73 | ~ spl102_270)), inference(forward_demodulation, [], [f3843, f1365])).
fof(f3843, plain, ((op2(e23, e21) = h4(e13)) | ~ spl102_270), inference(backward_demodulation, [], [f2627, f2664])).
fof(f2627, plain, (h4(e13) = op2(e23, h4(e11))), inference(backward_demodulation, [], [f919, f917])).
fof(f3949, plain, (~ spl102_118 | ~ spl102_304), inference(avatar_split_clause, [], [f3948, f2837, f1554])).
fof(f1554, plain, (spl102_118 <=> (e21 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl102_118])])).
fof(f3948, plain, (~ (e21 = op2(e20, e22)) | ~ spl102_304), inference(forward_demodulation, [], [f2601, f2838])).
fof(f2601, plain, ~ (op2(e20, e22) = h1(e11)), inference(backward_demodulation, [], [f399, f905])).
fof(f905, plain, (op2(e20, e20) = h1(e11)), inference(cnf_transformation, [], [f14])).
fof(f399, plain, ~ (op2(e20, e20) = op2(e20, e22)), inference(cnf_transformation, [], [f6])).
fof(f3944, plain, (~ spl102_120 | ~ spl102_124), inference(avatar_split_clause, [], [f3943, f1579, f1562])).
fof(f1562, plain, (spl102_120 <=> (e23 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl102_120])])).
fof(f3943, plain, (~ (e23 = op2(e20, e22)) | ~ spl102_124), inference(forward_demodulation, [], [f401, f1581])).
fof(f401, plain, ~ (op2(e20, e21) = op2(e20, e22)), inference(cnf_transformation, [], [f6])).
fof(f3942, plain, (spl102_290 | ~ spl102_107), inference(avatar_split_clause, [], [f3941, f1507, f2764])).
fof(f3941, plain, ((e22 = h2(e11)) | ~ spl102_107), inference(forward_demodulation, [], [f909, f1509])).
fof(f3931, plain, (~ spl102_44 | ~ spl102_60), inference(avatar_split_clause, [], [f3930, f1275, f1207])).
fof(f1207, plain, (spl102_44 <=> (e13 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl102_44])])).
fof(f3930, plain, (~ (e13 = op1(e11, e11)) | ~ spl102_60), inference(backward_demodulation, [], [f332, f1277])).
fof(f332, plain, ~ (op1(e10, e11) = op1(e11, e11)), inference(cnf_transformation, [], [f5])).
fof(f3929, plain, (~ spl102_43 | ~ spl102_39), inference(avatar_split_clause, [], [f3928, f1186, f1203])).
fof(f3928, plain, (~ (e12 = op1(e11, e11)) | ~ spl102_39), inference(forward_demodulation, [], [f359, f1188])).
fof(f3919, plain, (~ spl102_3 | ~ spl102_15), inference(avatar_split_clause, [], [f3918, f1084, f1033])).
fof(f1033, plain, (spl102_3 <=> (e12 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl102_3])])).
fof(f3918, plain, (~ (e12 = op1(e13, e13)) | ~ spl102_15), inference(forward_demodulation, [], [f370, f1086])).
fof(f370, plain, ~ (op1(e13, e10) = op1(e13, e13)), inference(cnf_transformation, [], [f5])).
fof(f3901, plain, (~ spl102_57 | ~ spl102_60), inference(avatar_contradiction_clause, [], [f3900])).
fof(f3900, plain, ($false | (~ spl102_57 | ~ spl102_60)), inference(subsumption_resolution, [], [f3899, f424])).
fof(f3899, plain, ((e10 = e13) | (~ spl102_57 | ~ spl102_60)), inference(backward_demodulation, [], [f1277, f1265])).
fof(f1265, plain, ((e10 = op1(e10, e11)) | ~ spl102_57), inference(avatar_component_clause, [], [f1263])).
fof(f1263, plain, (spl102_57 <=> (e10 = op1(e10, e11))), introduced(avatar_definition, [new_symbols(naming, [spl102_57])])).
fof(f3890, plain, (spl102_278 | ~ spl102_87), inference(avatar_split_clause, [], [f3889, f1422, f2703])).
fof(f3889, plain, ((e22 = h3(e11)) | ~ spl102_87), inference(backward_demodulation, [], [f913, f1424])).
fof(f3855, plain, (~ spl102_125 | ~ spl102_126), inference(avatar_contradiction_clause, [], [f3854])).
fof(f3854, plain, ($false | (~ spl102_125 | ~ spl102_126)), inference(subsumption_resolution, [], [f3853, f428])).
fof(f3853, plain, ((e20 = e21) | (~ spl102_125 | ~ spl102_126)), inference(backward_demodulation, [], [f1590, f1586])).
fof(f1586, plain, ((e20 = op2(e20, e20)) | ~ spl102_125), inference(avatar_component_clause, [], [f1584])).
fof(f3839, plain, (~ spl102_67 | ~ spl102_300), inference(avatar_split_clause, [], [f3835, f2816, f1337])).
fof(f1337, plain, (spl102_67 <=> (e22 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl102_67])])).
fof(f3835, plain, (~ (e22 = op2(e23, e23)) | ~ spl102_300), inference(backward_demodulation, [], [f2597, f2817])).
fof(f2597, plain, ~ (op2(e23, e23) = h1(e12)), inference(backward_demodulation, [], [f418, f2591])).
fof(f418, plain, ~ (op2(e23, e20) = op2(e23, e23)), inference(cnf_transformation, [], [f6])).
fof(f3838, plain, (~ spl102_95 | ~ spl102_300), inference(avatar_split_clause, [], [f3833, f2816, f1456])).
fof(f3833, plain, (~ (e22 = op2(e22, e20)) | ~ spl102_300), inference(backward_demodulation, [], [f2594, f2817])).
fof(f2594, plain, ~ (op2(e22, e20) = h1(e12)), inference(backward_demodulation, [], [f379, f2591])).
fof(f379, plain, ~ (op2(e22, e20) = op2(e23, e20)), inference(cnf_transformation, [], [f6])).
fof(f3832, plain, (~ spl102_60 | spl102_323), inference(avatar_contradiction_clause, [], [f3831])).
fof(f3831, plain, ($false | (~ spl102_60 | spl102_323)), inference(trivial_inequality_removal, [], [f3830])).
fof(f3830, plain, (~ (h4(e13) = h4(e13)) | (~ spl102_60 | spl102_323)), inference(forward_demodulation, [], [f2932, f1277])).
fof(f2932, plain, (~ (h4(e13) = h4(op1(e10, e11))) | spl102_323), inference(avatar_component_clause, [], [f2930])).
fof(f2930, plain, (spl102_323 <=> (h4(e13) = h4(op1(e10, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl102_323])])).
fof(f3808, plain, (~ spl102_110 | ~ spl102_304), inference(avatar_split_clause, [], [f3807, f2837, f1520])).
fof(f3807, plain, (~ (e21 = op2(e21, e20)) | ~ spl102_304), inference(forward_demodulation, [], [f2598, f2838])).
fof(f2598, plain, ~ (op2(e21, e20) = h1(e11)), inference(backward_demodulation, [], [f374, f905])).
fof(f374, plain, ~ (op2(e20, e20) = op2(e21, e20)), inference(cnf_transformation, [], [f6])).
fof(f3804, plain, (~ spl102_104 | ~ spl102_72), inference(avatar_split_clause, [], [f3803, f1358, f1494])).
fof(f3803, plain, (~ (e23 = op2(e21, e22)) | ~ spl102_72), inference(forward_demodulation, [], [f390, f1360])).
fof(f3785, plain, (~ spl102_56 | ~ spl102_60), inference(avatar_split_clause, [], [f3784, f1275, f1258])).
fof(f3784, plain, (~ (e13 = op1(e10, e12)) | ~ spl102_60), inference(forward_demodulation, [], [f353, f1277])).
fof(f353, plain, ~ (op1(e10, e11) = op1(e10, e12)), inference(cnf_transformation, [], [f5])).
fof(f3780, plain, (~ spl102_28 | ~ spl102_60), inference(avatar_split_clause, [], [f3779, f1275, f1139])).
fof(f1139, plain, (spl102_28 <=> (e13 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl102_28])])).
fof(f3779, plain, (~ (e13 = op1(e12, e11)) | ~ spl102_60), inference(forward_demodulation, [], [f333, f1277])).
fof(f333, plain, ~ (op1(e10, e11) = op1(e12, e11)), inference(cnf_transformation, [], [f5])).
fof(f3777, plain, (~ spl102_11 | ~ spl102_15), inference(avatar_split_clause, [], [f3776, f1084, f1067])).
fof(f1067, plain, (spl102_11 <=> (e12 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl102_11])])).
fof(f3776, plain, (~ (e12 = op1(e13, e11)) | ~ spl102_15), inference(forward_demodulation, [], [f368, f1086])).
fof(f368, plain, ~ (op1(e13, e10) = op1(e13, e11)), inference(cnf_transformation, [], [f5])).
fof(f3775, plain, (~ spl102_12 | ~ spl102_60), inference(avatar_split_clause, [], [f3774, f1275, f1071])).
fof(f3774, plain, (~ (e13 = op1(e13, e11)) | ~ spl102_60), inference(forward_demodulation, [], [f334, f1277])).
fof(f334, plain, ~ (op1(e10, e11) = op1(e13, e11)), inference(cnf_transformation, [], [f5])).
fof(f3753, plain, (~ spl102_41 | ~ spl102_42), inference(avatar_contradiction_clause, [], [f3752])).
fof(f3752, plain, ($false | (~ spl102_41 | ~ spl102_42)), inference(subsumption_resolution, [], [f3751, f422])).
fof(f3751, plain, ((e10 = e11) | (~ spl102_41 | ~ spl102_42)), inference(backward_demodulation, [], [f1201, f1197])).
fof(f1201, plain, ((e11 = op1(e11, e11)) | ~ spl102_42), inference(avatar_component_clause, [], [f1199])).
fof(f3750, plain, (~ spl102_40 | ~ spl102_48), inference(avatar_split_clause, [], [f3748, f1224, f1190])).
fof(f3748, plain, (~ (e13 = op1(e11, e12)) | ~ spl102_48), inference(backward_demodulation, [], [f357, f1226])).
fof(f357, plain, ~ (op1(e11, e10) = op1(e11, e12)), inference(cnf_transformation, [], [f5])).
fof(f3743, plain, (~ spl102_86 | spl102_282), inference(avatar_contradiction_clause, [], [f3742])).
fof(f3742, plain, ($false | (~ spl102_86 | spl102_282)), inference(subsumption_resolution, [], [f3741, f2726])).
fof(f2726, plain, (~ (e21 = h3(e11)) | spl102_282), inference(avatar_component_clause, [], [f2724])).
fof(f3741, plain, ((e21 = h3(e11)) | ~ spl102_86), inference(backward_demodulation, [], [f913, f1420])).
fof(f3731, plain, (spl102_364 | ~ spl102_108), inference(avatar_split_clause, [], [f3730, f1511, f3195])).
fof(f1511, plain, (spl102_108 <=> (e23 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl102_108])])).
fof(f3730, plain, ((e23 = h2(e11)) | ~ spl102_108), inference(backward_demodulation, [], [f909, f1513])).
fof(f1513, plain, ((e23 = op2(e21, e21)) | ~ spl102_108), inference(avatar_component_clause, [], [f1511])).
fof(f3729, plain, (~ spl102_109 | ~ spl102_112), inference(avatar_contradiction_clause, [], [f3728])).
fof(f3728, plain, ($false | (~ spl102_109 | ~ spl102_112)), inference(subsumption_resolution, [], [f3727, f430])).
fof(f3727, plain, ((e20 = e23) | (~ spl102_109 | ~ spl102_112)), inference(forward_demodulation, [], [f1530, f1518])).
fof(f3722, plain, (~ spl102_117 | ~ spl102_119), inference(avatar_contradiction_clause, [], [f3721])).
fof(f3721, plain, ($false | (~ spl102_117 | ~ spl102_119)), inference(subsumption_resolution, [], [f3720, f429])).
fof(f3720, plain, ((e20 = e22) | (~ spl102_117 | ~ spl102_119)), inference(forward_demodulation, [], [f1560, f1552])).
fof(f3703, plain, (~ spl102_74 | spl102_106 | ~ spl102_215), inference(avatar_contradiction_clause, [], [f3702])).
fof(f3702, plain, ($false | (~ spl102_74 | spl102_106 | ~ spl102_215)), inference(subsumption_resolution, [], [f3701, f1504])).
fof(f3701, plain, ((e21 = op2(e21, e21)) | (~ spl102_74 | ~ spl102_215)), inference(forward_demodulation, [], [f2264, f1369])).
fof(f3696, plain, (spl102_124 | ~ spl102_304), inference(avatar_split_clause, [], [f3691, f2837, f1579])).
fof(f3691, plain, ((e23 = op2(e20, e21)) | ~ spl102_304), inference(backward_demodulation, [], [f2603, f2838])).
fof(f2603, plain, (e23 = op2(e20, h1(e11))), inference(backward_demodulation, [], [f903, f905])).
fof(f3695, plain, (~ spl102_114 | ~ spl102_304), inference(avatar_split_clause, [], [f3690, f2837, f1537])).
fof(f1537, plain, (spl102_114 <=> (e21 = op2(e20, e23))), introduced(avatar_definition, [new_symbols(naming, [spl102_114])])).
fof(f3690, plain, (~ (e21 = op2(e20, e23)) | ~ spl102_304), inference(backward_demodulation, [], [f2602, f2838])).
fof(f2602, plain, ~ (op2(e20, e23) = h1(e11)), inference(backward_demodulation, [], [f400, f905])).
fof(f400, plain, ~ (op2(e20, e20) = op2(e20, e23)), inference(cnf_transformation, [], [f6])).
fof(f3694, plain, (~ spl102_122 | ~ spl102_304), inference(avatar_split_clause, [], [f3689, f2837, f1571])).
fof(f1571, plain, (spl102_122 <=> (e21 = op2(e20, e21))), introduced(avatar_definition, [new_symbols(naming, [spl102_122])])).
fof(f3689, plain, (~ (e21 = op2(e20, e21)) | ~ spl102_304), inference(backward_demodulation, [], [f2600, f2838])).
fof(f2600, plain, ~ (op2(e20, e21) = h1(e11)), inference(backward_demodulation, [], [f398, f905])).
fof(f398, plain, ~ (op2(e20, e20) = op2(e20, e21)), inference(cnf_transformation, [], [f6])).
fof(f3693, plain, (~ spl102_94 | ~ spl102_304), inference(avatar_split_clause, [], [f3688, f2837, f1452])).
fof(f1452, plain, (spl102_94 <=> (e21 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl102_94])])).
fof(f3688, plain, (~ (e21 = op2(e22, e20)) | ~ spl102_304), inference(backward_demodulation, [], [f2599, f2838])).
fof(f2599, plain, ~ (op2(e22, e20) = h1(e11)), inference(backward_demodulation, [], [f375, f905])).
fof(f375, plain, ~ (op2(e20, e20) = op2(e22, e20)), inference(cnf_transformation, [], [f6])).
fof(f3671, plain, (~ spl102_115 | ~ spl102_99), inference(avatar_split_clause, [], [f3670, f1473, f1541])).
fof(f3670, plain, (~ (e22 = op2(e20, e23)) | ~ spl102_99), inference(forward_demodulation, [], [f392, f1475])).
fof(f1475, plain, ((e22 = op2(e21, e23)) | ~ spl102_99), inference(avatar_component_clause, [], [f1473])).
fof(f3656, plain, (~ spl102_278 | ~ spl102_91), inference(avatar_split_clause, [], [f3655, f1439, f2703])).
fof(f3655, plain, (~ (e22 = h3(e11)) | ~ spl102_91), inference(forward_demodulation, [], [f2618, f1441])).
fof(f3648, plain, (spl102_300 | ~ spl102_79), inference(avatar_split_clause, [], [f3647, f1388, f2816])).
fof(f3647, plain, ((e22 = h1(e12)) | ~ spl102_79), inference(forward_demodulation, [], [f2591, f1390])).
fof(f3646, plain, (spl102_274 | ~ spl102_65), inference(avatar_split_clause, [], [f3645, f1329, f2683])).
fof(f3645, plain, ((e20 = h4(e11)) | ~ spl102_65), inference(forward_demodulation, [], [f917, f1331])).
fof(f3622, plain, (~ spl102_47 | ~ spl102_15), inference(avatar_split_clause, [], [f3621, f1084, f1220])).
fof(f1220, plain, (spl102_47 <=> (e12 = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl102_47])])).
fof(f3621, plain, (~ (e12 = op1(e11, e10)) | ~ spl102_15), inference(forward_demodulation, [], [f330, f1086])).
fof(f330, plain, ~ (op1(e11, e10) = op1(e13, e10)), inference(cnf_transformation, [], [f5])).
fof(f3612, plain, (~ spl102_24 | ~ spl102_20), inference(avatar_split_clause, [], [f3611, f1105, f1122])).
fof(f3611, plain, (~ (e13 = op1(e12, e12)) | ~ spl102_20), inference(forward_demodulation, [], [f367, f1107])).
fof(f3606, plain, (~ spl102_15 | spl102_318), inference(avatar_contradiction_clause, [], [f3605])).
fof(f3605, plain, ($false | (~ spl102_15 | spl102_318)), inference(trivial_inequality_removal, [], [f3604])).
fof(f3604, plain, (~ (h4(e12) = h4(e12)) | (~ spl102_15 | spl102_318)), inference(forward_demodulation, [], [f2912, f1086])).
fof(f2912, plain, (~ (h4(e12) = h4(op1(e13, e10))) | spl102_318), inference(avatar_component_clause, [], [f2910])).
fof(f2910, plain, (spl102_318 <=> (h4(e12) = h4(op1(e13, e10)))), introduced(avatar_definition, [new_symbols(naming, [spl102_318])])).
fof(f3597, plain, (~ spl102_13 | ~ spl102_15), inference(avatar_contradiction_clause, [], [f3596])).
fof(f3596, plain, ($false | (~ spl102_13 | ~ spl102_15)), inference(subsumption_resolution, [], [f3595, f423])).
fof(f3595, plain, ((e10 = e12) | (~ spl102_13 | ~ spl102_15)), inference(backward_demodulation, [], [f1086, f1078])).
fof(f1078, plain, ((e10 = op1(e13, e10)) | ~ spl102_13), inference(avatar_component_clause, [], [f1076])).
fof(f1076, plain, (spl102_13 <=> (e10 = op1(e13, e10))), introduced(avatar_definition, [new_symbols(naming, [spl102_13])])).
fof(f3594, plain, (~ spl102_15 | ~ spl102_16), inference(avatar_contradiction_clause, [], [f3593])).
fof(f3593, plain, ($false | (~ spl102_15 | ~ spl102_16)), inference(subsumption_resolution, [], [f3592, f427])).
fof(f427, plain, ~ (e12 = e13), inference(cnf_transformation, [], [f7])).
fof(f3592, plain, ((e12 = e13) | (~ spl102_15 | ~ spl102_16)), inference(backward_demodulation, [], [f1090, f1086])).
fof(f1090, plain, ((e13 = op1(e13, e10)) | ~ spl102_16), inference(avatar_component_clause, [], [f1088])).
fof(f1088, plain, (spl102_16 <=> (e13 = op1(e13, e10))), introduced(avatar_definition, [new_symbols(naming, [spl102_16])])).
fof(f3586, plain, (~ spl102_17 | ~ spl102_20), inference(avatar_contradiction_clause, [], [f3585])).
fof(f3585, plain, ($false | (~ spl102_17 | ~ spl102_20)), inference(subsumption_resolution, [], [f3584, f424])).
fof(f3584, plain, ((e10 = e13) | (~ spl102_17 | ~ spl102_20)), inference(backward_demodulation, [], [f1107, f1095])).
fof(f3578, plain, (~ spl102_4 | ~ spl102_20), inference(avatar_split_clause, [], [f3573, f1105, f1037])).
fof(f3573, plain, (~ (e13 = op1(e13, e13)) | ~ spl102_20), inference(backward_demodulation, [], [f349, f1107])).
fof(f3568, plain, (~ spl102_26 | ~ spl102_27), inference(avatar_contradiction_clause, [], [f3567])).
fof(f3567, plain, ($false | (~ spl102_26 | ~ spl102_27)), inference(subsumption_resolution, [], [f3566, f425])).
fof(f3566, plain, ((e11 = e12) | (~ spl102_26 | ~ spl102_27)), inference(backward_demodulation, [], [f1137, f1133])).
fof(f1137, plain, ((e12 = op1(e12, e11)) | ~ spl102_27), inference(avatar_component_clause, [], [f1135])).
fof(f3565, plain, (~ spl102_19 | ~ spl102_27), inference(avatar_split_clause, [], [f3562, f1135, f1101])).
fof(f3562, plain, (~ (e12 = op1(e12, e13)) | ~ spl102_27), inference(backward_demodulation, [], [f366, f1137])).
fof(f3564, plain, (~ spl102_23 | ~ spl102_27), inference(avatar_split_clause, [], [f3561, f1135, f1118])).
fof(f3561, plain, (~ (e12 = op1(e12, e12)) | ~ spl102_27), inference(backward_demodulation, [], [f365, f1137])).
fof(f365, plain, ~ (op1(e12, e11) = op1(e12, e12)), inference(cnf_transformation, [], [f5])).
fof(f3558, plain, (~ spl102_21 | ~ spl102_29), inference(avatar_split_clause, [], [f3554, f1144, f1110])).
fof(f3554, plain, (~ (e10 = op1(e12, e12)) | ~ spl102_29), inference(backward_demodulation, [], [f363, f1146])).
fof(f3557, plain, (~ spl102_25 | ~ spl102_29), inference(avatar_split_clause, [], [f3553, f1144, f1127])).
fof(f3553, plain, (~ (e10 = op1(e12, e11)) | ~ spl102_29), inference(backward_demodulation, [], [f362, f1146])).
fof(f362, plain, ~ (op1(e12, e10) = op1(e12, e11)), inference(cnf_transformation, [], [f5])).
fof(f3538, plain, (~ spl102_34 | ~ spl102_42), inference(avatar_split_clause, [], [f3534, f1199, f1165])).
fof(f3534, plain, (~ (e11 = op1(e11, e13)) | ~ spl102_42), inference(backward_demodulation, [], [f360, f1201])).
fof(f3536, plain, (~ spl102_10 | ~ spl102_42), inference(avatar_split_clause, [], [f3532, f1199, f1063])).
fof(f3532, plain, (~ (e11 = op1(e13, e11)) | ~ spl102_42), inference(backward_demodulation, [], [f336, f1201])).
fof(f3529, plain, (~ spl102_37 | ~ spl102_45), inference(avatar_split_clause, [], [f3524, f1212, f1178])).
fof(f3524, plain, (~ (e10 = op1(e11, e12)) | ~ spl102_45), inference(backward_demodulation, [], [f357, f1214])).
fof(f3528, plain, (~ spl102_41 | ~ spl102_45), inference(avatar_split_clause, [], [f3523, f1212, f1195])).
fof(f3523, plain, (~ (e10 = op1(e11, e11)) | ~ spl102_45), inference(backward_demodulation, [], [f356, f1214])).
fof(f356, plain, ~ (op1(e11, e10) = op1(e11, e11)), inference(cnf_transformation, [], [f5])).
fof(f3505, plain, (~ spl102_5 | ~ spl102_53), inference(avatar_split_clause, [], [f3501, f1246, f1042])).
fof(f3501, plain, (~ (e10 = op1(e13, e12)) | ~ spl102_53), inference(backward_demodulation, [], [f340, f1248])).
fof(f340, plain, ~ (op1(e10, e12) = op1(e13, e12)), inference(cnf_transformation, [], [f5])).
fof(f3503, plain, (~ spl102_37 | ~ spl102_53), inference(avatar_split_clause, [], [f3499, f1246, f1178])).
fof(f3499, plain, (~ (e10 = op1(e11, e12)) | ~ spl102_53), inference(backward_demodulation, [], [f338, f1248])).
fof(f3488, plain, (spl102_60 | ~ spl102_62), inference(avatar_split_clause, [], [f3481, f1284, f1275])).
fof(f3481, plain, ((e13 = op1(e10, e11)) | ~ spl102_62), inference(backward_demodulation, [], [f900, f1286])).
fof(f900, plain, (e13 = op1(e10, op1(e10, e10))), inference(cnf_transformation, [], [f12])).
fof(f12, plain, ((e13 = op1(e10, op1(e10, e10))) & (e12 = op1(op1(e10, op1(e10, e10)), e10)) & (op1(e10, e10) = e11)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG109+1.p', ax12)).
fof(f3487, plain, (~ spl102_50 | ~ spl102_62), inference(avatar_split_clause, [], [f3480, f1284, f1233])).
fof(f1233, plain, (spl102_50 <=> (e11 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl102_50])])).
fof(f3480, plain, (~ (e11 = op1(e10, e13)) | ~ spl102_62), inference(backward_demodulation, [], [f352, f1286])).
fof(f352, plain, ~ (op1(e10, e10) = op1(e10, e13)), inference(cnf_transformation, [], [f5])).
fof(f3486, plain, (~ spl102_54 | ~ spl102_62), inference(avatar_split_clause, [], [f3479, f1284, f1250])).
fof(f1250, plain, (spl102_54 <=> (e11 = op1(e10, e12))), introduced(avatar_definition, [new_symbols(naming, [spl102_54])])).
fof(f3479, plain, (~ (e11 = op1(e10, e12)) | ~ spl102_62), inference(backward_demodulation, [], [f351, f1286])).
fof(f351, plain, ~ (op1(e10, e10) = op1(e10, e12)), inference(cnf_transformation, [], [f5])).
fof(f3485, plain, (~ spl102_58 | ~ spl102_62), inference(avatar_split_clause, [], [f3478, f1284, f1267])).
fof(f1267, plain, (spl102_58 <=> (e11 = op1(e10, e11))), introduced(avatar_definition, [new_symbols(naming, [spl102_58])])).
fof(f3478, plain, (~ (e11 = op1(e10, e11)) | ~ spl102_62), inference(backward_demodulation, [], [f350, f1286])).
fof(f350, plain, ~ (op1(e10, e10) = op1(e10, e11)), inference(cnf_transformation, [], [f5])).
fof(f3484, plain, (~ spl102_14 | ~ spl102_62), inference(avatar_split_clause, [], [f3477, f1284, f1080])).
fof(f1080, plain, (spl102_14 <=> (e11 = op1(e13, e10))), introduced(avatar_definition, [new_symbols(naming, [spl102_14])])).
fof(f3477, plain, (~ (e11 = op1(e13, e10)) | ~ spl102_62), inference(backward_demodulation, [], [f328, f1286])).
fof(f328, plain, ~ (op1(e10, e10) = op1(e13, e10)), inference(cnf_transformation, [], [f5])).
fof(f3483, plain, (~ spl102_30 | ~ spl102_62), inference(avatar_split_clause, [], [f3476, f1284, f1148])).
fof(f1148, plain, (spl102_30 <=> (e11 = op1(e12, e10))), introduced(avatar_definition, [new_symbols(naming, [spl102_30])])).
fof(f3476, plain, (~ (e11 = op1(e12, e10)) | ~ spl102_62), inference(backward_demodulation, [], [f327, f1286])).
fof(f327, plain, ~ (op1(e10, e10) = op1(e12, e10)), inference(cnf_transformation, [], [f5])).
fof(f3482, plain, (~ spl102_46 | ~ spl102_62), inference(avatar_split_clause, [], [f3475, f1284, f1216])).
fof(f1216, plain, (spl102_46 <=> (e11 = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl102_46])])).
fof(f3475, plain, (~ (e11 = op1(e11, e10)) | ~ spl102_62), inference(backward_demodulation, [], [f326, f1286])).
fof(f326, plain, ~ (op1(e10, e10) = op1(e11, e10)), inference(cnf_transformation, [], [f5])).
fof(f3474, plain, (~ spl102_65 | ~ spl102_66), inference(avatar_contradiction_clause, [], [f3473])).
fof(f3473, plain, ($false | (~ spl102_65 | ~ spl102_66)), inference(subsumption_resolution, [], [f3472, f428])).
fof(f3472, plain, ((e20 = e21) | (~ spl102_65 | ~ spl102_66)), inference(backward_demodulation, [], [f1335, f1331])).
fof(f1335, plain, ((e21 = op2(e23, e23)) | ~ spl102_66), inference(avatar_component_clause, [], [f1333])).
fof(f1333, plain, (spl102_66 <=> (e21 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl102_66])])).
fof(f3470, plain, (spl102_270 | ~ spl102_66), inference(avatar_split_clause, [], [f3467, f1333, f2663])).
fof(f3467, plain, ((e21 = h4(e11)) | ~ spl102_66), inference(backward_demodulation, [], [f917, f1335])).
fof(f3456, plain, (~ spl102_294 | ~ spl102_74), inference(avatar_split_clause, [], [f3453, f1367, f2784])).
fof(f3453, plain, (~ (e21 = h2(e11)) | ~ spl102_74), inference(backward_demodulation, [], [f2608, f1369])).
fof(f3449, plain, (~ spl102_79 | ~ spl102_80), inference(avatar_contradiction_clause, [], [f3448])).
fof(f3448, plain, ($false | (~ spl102_79 | ~ spl102_80)), inference(subsumption_resolution, [], [f3447, f433])).
fof(f3447, plain, ((e22 = e23) | (~ spl102_79 | ~ spl102_80)), inference(backward_demodulation, [], [f1394, f1390])).
fof(f1394, plain, ((e23 = op2(e23, e20)) | ~ spl102_80), inference(avatar_component_clause, [], [f1392])).
fof(f1392, plain, (spl102_80 <=> (e23 = op2(e23, e20))), introduced(avatar_definition, [new_symbols(naming, [spl102_80])])).
fof(f3430, plain, (spl102_286 | ~ spl102_85), inference(avatar_split_clause, [], [f3429, f1414, f2744])).
fof(f2744, plain, (spl102_286 <=> (e20 = h3(e11))), introduced(avatar_definition, [new_symbols(naming, [spl102_286])])).
fof(f1414, plain, (spl102_85 <=> (e20 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl102_85])])).
fof(f3429, plain, ((e20 = h3(e11)) | ~ spl102_85), inference(backward_demodulation, [], [f913, f1416])).
fof(f1416, plain, ((e20 = op2(e22, e22)) | ~ spl102_85), inference(avatar_component_clause, [], [f1414])).
fof(f3415, plain, (~ spl102_89 | ~ spl102_93), inference(avatar_split_clause, [], [f3410, f1448, f1431])).
fof(f3410, plain, (~ (e20 = op2(e22, e21)) | ~ spl102_93), inference(backward_demodulation, [], [f410, f1450])).
fof(f1450, plain, ((e20 = op2(e22, e20)) | ~ spl102_93), inference(avatar_component_clause, [], [f1448])).
fof(f410, plain, ~ (op2(e22, e20) = op2(e22, e21)), inference(cnf_transformation, [], [f6])).
fof(f3394, plain, (~ spl102_294 | ~ spl102_102), inference(avatar_split_clause, [], [f3390, f1486, f2784])).
fof(f3390, plain, (~ (e21 = h2(e11)) | ~ spl102_102), inference(backward_demodulation, [], [f2610, f1488])).
fof(f3393, plain, (~ spl102_282 | ~ spl102_102), inference(avatar_split_clause, [], [f3389, f1486, f2724])).
fof(f3389, plain, (~ (e21 = h3(e11)) | ~ spl102_102), inference(backward_demodulation, [], [f2615, f1488])).
fof(f3388, plain, (spl102_298 | ~ spl102_105), inference(avatar_split_clause, [], [f3387, f1499, f2805])).
fof(f1499, plain, (spl102_105 <=> (e20 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl102_105])])).
fof(f3387, plain, ((e20 = h2(e11)) | ~ spl102_105), inference(backward_demodulation, [], [f909, f1501])).
fof(f1501, plain, ((e20 = op2(e21, e21)) | ~ spl102_105), inference(avatar_component_clause, [], [f1499])).
fof(f3382, plain, (~ spl102_97 | ~ spl102_109), inference(avatar_split_clause, [], [f3375, f1516, f1465])).
fof(f3375, plain, (~ (e20 = op2(e21, e23)) | ~ spl102_109), inference(backward_demodulation, [], [f406, f1518])).
fof(f3381, plain, (~ spl102_101 | ~ spl102_109), inference(avatar_split_clause, [], [f3374, f1516, f1482])).
fof(f3374, plain, (~ (e20 = op2(e21, e22)) | ~ spl102_109), inference(backward_demodulation, [], [f405, f1518])).
fof(f3358, plain, (~ spl102_101 | ~ spl102_117), inference(avatar_split_clause, [], [f3353, f1550, f1482])).
fof(f3353, plain, (~ (e20 = op2(e21, e22)) | ~ spl102_117), inference(backward_demodulation, [], [f386, f1552])).
fof(f3356, plain, (~ spl102_286 | ~ spl102_117), inference(avatar_split_clause, [], [f3351, f1550, f2744])).
fof(f3351, plain, (~ (e20 = h3(e11)) | ~ spl102_117), inference(backward_demodulation, [], [f2614, f1552])).
fof(f3338, plain, (spl102_304 | ~ spl102_126), inference(avatar_split_clause, [], [f3337, f1588, f2837])).
fof(f3337, plain, ((e21 = h1(e11)) | ~ spl102_126), inference(backward_demodulation, [], [f905, f1590])).
fof(f2979, plain, (~ spl102_308 | ~ spl102_309 | ~ spl102_310 | ~ spl102_311 | ~ spl102_312 | ~ spl102_313 | ~ spl102_314 | ~ spl102_315 | ~ spl102_316 | spl102_271 | spl102_267 | spl102_263 | ~ spl102_318 | ~ spl102_319 | ~ spl102_320 | ~ spl102_321 | ~ spl102_322 | ~ spl102_323 | ~ spl102_324), inference(avatar_split_clause, [], [f2978, f2934, f2930, f2926, f2922, f2918, f2914, f2910, f2629, f2649, f2669, f2902, f2898, f2894, f2890, f2886, f2882, f2878, f2874, f2870])).
fof(f2669, plain, (spl102_271 <=> sP99), introduced(avatar_definition, [new_symbols(naming, [spl102_271])])).
fof(f2649, plain, (spl102_267 <=> sP100), introduced(avatar_definition, [new_symbols(naming, [spl102_267])])).
fof(f2629, plain, (spl102_263 <=> sP101), introduced(avatar_definition, [new_symbols(naming, [spl102_263])])).
fof(f2978, plain, (~ (h4(e11) = h4(op1(e10, e10))) | ~ (h4(e13) = h4(op1(e10, e11))) | ~ (h4(op1(e10, e12)) = op2(e23, h4(e12))) | ~ (h4(op1(e10, e13)) = op2(e23, h4(e13))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), e23)) | ~ (h4(op1(e12, e10)) = op2(h4(e12), e23)) | ~ (h4(e12) = h4(op1(e13, e10))) | sP101 | sP100 | sP99 | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11)))), inference(forward_demodulation, [], [f2977, f917])).
fof(f2977, plain, (~ (op2(e23, e23) = h4(op1(e10, e10))) | ~ (h4(e13) = h4(op1(e10, e11))) | ~ (h4(op1(e10, e12)) = op2(e23, h4(e12))) | ~ (h4(op1(e10, e13)) = op2(e23, h4(e13))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), e23)) | ~ (h4(op1(e12, e10)) = op2(h4(e12), e23)) | ~ (h4(e12) = h4(op1(e13, e10))) | sP101 | sP100 | sP99 | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11)))), inference(forward_demodulation, [], [f2976, f916])).
fof(f2976, plain, (~ (h4(e13) = h4(op1(e10, e11))) | ~ (h4(op1(e10, e12)) = op2(e23, h4(e12))) | ~ (h4(op1(e10, e13)) = op2(e23, h4(e13))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), e23)) | ~ (h4(op1(e12, e10)) = op2(h4(e12), e23)) | ~ (h4(e12) = h4(op1(e13, e10))) | sP101 | sP100 | sP99 | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f2975, f2627])).
fof(f2975, plain, (~ (h4(op1(e10, e11)) = op2(e23, h4(e11))) | ~ (h4(op1(e10, e12)) = op2(e23, h4(e12))) | ~ (h4(op1(e10, e13)) = op2(e23, h4(e13))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), e23)) | ~ (h4(op1(e12, e10)) = op2(h4(e12), e23)) | ~ (h4(e12) = h4(op1(e13, e10))) | sP101 | sP100 | sP99 | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f2974, f916])).
fof(f2974, plain, (~ (h4(op1(e10, e12)) = op2(e23, h4(e12))) | ~ (h4(op1(e10, e13)) = op2(e23, h4(e13))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), e23)) | ~ (h4(op1(e12, e10)) = op2(h4(e12), e23)) | ~ (h4(e12) = h4(op1(e13, e10))) | sP101 | sP100 | sP99 | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f2973, f916])).
fof(f2973, plain, (~ (h4(op1(e10, e13)) = op2(e23, h4(e13))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), e23)) | ~ (h4(op1(e12, e10)) = op2(h4(e12), e23)) | ~ (h4(e12) = h4(op1(e13, e10))) | sP101 | sP100 | sP99 | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f2972, f916])).
fof(f2972, plain, (~ (h4(op1(e11, e10)) = op2(h4(e11), e23)) | ~ (h4(op1(e12, e10)) = op2(h4(e12), e23)) | ~ (h4(e12) = h4(op1(e13, e10))) | sP101 | sP100 | sP99 | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f2971, f916])).
fof(f2971, plain, (~ (h4(op1(e12, e10)) = op2(h4(e12), e23)) | ~ (h4(e12) = h4(op1(e13, e10))) | sP101 | sP100 | sP99 | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f2970, f916])).
fof(f2970, plain, (~ (h4(e12) = h4(op1(e13, e10))) | sP101 | sP100 | sP99 | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f2969, f2621])).
fof(f2969, plain, (~ (h4(op1(e13, e10)) = op2(h4(e13), e23)) | sP101 | sP100 | sP99 | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f2968, f916])).
fof(f2968, plain, (sP101 | sP100 | sP99 | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(subsumption_resolution, [], [f980, f916])).
fof(f980, plain, (~ (e23 = h4(e10)) | sP101 | sP100 | sP99 | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(cnf_transformation, [], [f127])).
fof(f127, plain, (((~ (e23 = h4(e13)) & ~ (e23 = h4(e12)) & ~ (e23 = h4(e11)) & ~ (e23 = h4(e10))) | sP101 | sP100 | sP99 | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) & ((~ (e23 = h3(e13)) & ~ (e23 = h3(e12)) & ~ (e23 = h3(e11)) & ~ (e23 = h3(e10))) | sP98 | sP97 | sP96 | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) & ((~ (e23 = h2(e13)) & ~ (e23 = h2(e12)) & ~ (e23 = h2(e11)) & ~ (e23 = h2(e10))) | sP95 | sP94 | sP93 | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) & ((~ (e23 = h1(e13)) & ~ (e23 = h1(e12)) & ~ (e23 = h1(e11)) & ~ (e23 = h1(e10))) | sP92 | sP91 | sP90 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(definition_folding, [], [f22, e126, e125, e124, e123, e122, e121, e120, e119, e118, e117, e116, e115])).
fof(f115, plain, ((~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10))) | ~ sP90), inference(usedef, [], [e115])).
fof(e115, plain, (sP90 <=> (~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP90])])).
fof(f116, plain, ((~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10))) | ~ sP91), inference(usedef, [], [e116])).
fof(e116, plain, (sP91 <=> (~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP91])])).
fof(f117, plain, ((~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10))) | ~ sP92), inference(usedef, [], [e117])).
fof(e117, plain, (sP92 <=> (~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP92])])).
fof(f118, plain, ((~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10))) | ~ sP93), inference(usedef, [], [e118])).
fof(e118, plain, (sP93 <=> (~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP93])])).
fof(f119, plain, ((~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10))) | ~ sP94), inference(usedef, [], [e119])).
fof(e119, plain, (sP94 <=> (~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP94])])).
fof(f120, plain, ((~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10))) | ~ sP95), inference(usedef, [], [e120])).
fof(e120, plain, (sP95 <=> (~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP95])])).
fof(f121, plain, ((~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10))) | ~ sP96), inference(usedef, [], [e121])).
fof(e121, plain, (sP96 <=> (~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP96])])).
fof(f122, plain, ((~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10))) | ~ sP97), inference(usedef, [], [e122])).
fof(e122, plain, (sP97 <=> (~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP97])])).
fof(f123, plain, ((~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10))) | ~ sP98), inference(usedef, [], [e123])).
fof(e123, plain, (sP98 <=> (~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP98])])).
fof(f124, plain, ((~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10))) | ~ sP99), inference(usedef, [], [e124])).
fof(e124, plain, (sP99 <=> (~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP99])])).
fof(f125, plain, ((~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10))) | ~ sP100), inference(usedef, [], [e125])).
fof(e125, plain, (sP100 <=> (~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP100])])).
fof(f126, plain, ((~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10))) | ~ sP101), inference(usedef, [], [e126])).
fof(e126, plain, (sP101 <=> (~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP101])])).
fof(f22, plain, (((~ (e23 = h4(e13)) & ~ (e23 = h4(e12)) & ~ (e23 = h4(e11)) & ~ (e23 = h4(e10))) | (~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10))) | (~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10))) | (~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10))) | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) & ((~ (e23 = h3(e13)) & ~ (e23 = h3(e12)) & ~ (e23 = h3(e11)) & ~ (e23 = h3(e10))) | (~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10))) | (~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10))) | (~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10))) | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) & ((~ (e23 = h2(e13)) & ~ (e23 = h2(e12)) & ~ (e23 = h2(e11)) & ~ (e23 = h2(e10))) | (~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10))) | (~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10))) | (~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) & ((~ (e23 = h1(e13)) & ~ (e23 = h1(e12)) & ~ (e23 = h1(e11)) & ~ (e23 = h1(e10))) | (~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10))) | (~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10))) | (~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10))) | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(ennf_transformation, [], [f19])).
fof(f19, plain, ~ ((((e23 = h4(e13)) | (e23 = h4(e12)) | (e23 = h4(e11)) | (e23 = h4(e10))) & ((e22 = h4(e13)) | (e22 = h4(e12)) | (e22 = h4(e11)) | (e22 = h4(e10))) & ((e21 = h4(e13)) | (e21 = h4(e12)) | (e21 = h4(e11)) | (e21 = h4(e10))) & ((e20 = h4(e13)) | (e20 = h4(e12)) | (e20 = h4(e11)) | (e20 = h4(e10))) & (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) & (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) & (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) & (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) & (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) & (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) & (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) & (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) & (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) & (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) & (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) & (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) & (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) & (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) & (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) & (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) | (((e23 = h3(e13)) | (e23 = h3(e12)) | (e23 = h3(e11)) | (e23 = h3(e10))) & ((e22 = h3(e13)) | (e22 = h3(e12)) | (e22 = h3(e11)) | (e22 = h3(e10))) & ((e21 = h3(e13)) | (e21 = h3(e12)) | (e21 = h3(e11)) | (e21 = h3(e10))) & ((e20 = h3(e13)) | (e20 = h3(e12)) | (e20 = h3(e11)) | (e20 = h3(e10))) & (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) & (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) & (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) & (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) & (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) & (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) & (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) & (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) & (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) & (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) & (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) & (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) & (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) & (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) & (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) & (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) | (((e23 = h2(e13)) | (e23 = h2(e12)) | (e23 = h2(e11)) | (e23 = h2(e10))) & ((e22 = h2(e13)) | (e22 = h2(e12)) | (e22 = h2(e11)) | (e22 = h2(e10))) & ((e21 = h2(e13)) | (e21 = h2(e12)) | (e21 = h2(e11)) | (e21 = h2(e10))) & ((e20 = h2(e13)) | (e20 = h2(e12)) | (e20 = h2(e11)) | (e20 = h2(e10))) & (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) & (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) & (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) & (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) & (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) & (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) & (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) & (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) & (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) & (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) & (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) & (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) & (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) & (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) & (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) & (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) | (((e23 = h1(e13)) | (e23 = h1(e12)) | (e23 = h1(e11)) | (e23 = h1(e10))) & ((e22 = h1(e13)) | (e22 = h1(e12)) | (e22 = h1(e11)) | (e22 = h1(e10))) & ((e21 = h1(e13)) | (e21 = h1(e12)) | (e21 = h1(e11)) | (e21 = h1(e10))) & ((e20 = h1(e13)) | (e20 = h1(e12)) | (e20 = h1(e11)) | (e20 = h1(e10))) & (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) & (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) & (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) & (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) & (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) & (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) & (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) & (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) & (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) & (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) & (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) & (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) & (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) & (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) & (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) & (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(negated_conjecture, [], [f18])).
fof(f18, plain, ~ ((((e23 = h4(e13)) | (e23 = h4(e12)) | (e23 = h4(e11)) | (e23 = h4(e10))) & ((e22 = h4(e13)) | (e22 = h4(e12)) | (e22 = h4(e11)) | (e22 = h4(e10))) & ((e21 = h4(e13)) | (e21 = h4(e12)) | (e21 = h4(e11)) | (e21 = h4(e10))) & ((e20 = h4(e13)) | (e20 = h4(e12)) | (e20 = h4(e11)) | (e20 = h4(e10))) & (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) & (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) & (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) & (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) & (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) & (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) & (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) & (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) & (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) & (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) & (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) & (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) & (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) & (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) & (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) & (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) | (((e23 = h3(e13)) | (e23 = h3(e12)) | (e23 = h3(e11)) | (e23 = h3(e10))) & ((e22 = h3(e13)) | (e22 = h3(e12)) | (e22 = h3(e11)) | (e22 = h3(e10))) & ((e21 = h3(e13)) | (e21 = h3(e12)) | (e21 = h3(e11)) | (e21 = h3(e10))) & ((e20 = h3(e13)) | (e20 = h3(e12)) | (e20 = h3(e11)) | (e20 = h3(e10))) & (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) & (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) & (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) & (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) & (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) & (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) & (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) & (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) & (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) & (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) & (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) & (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) & (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) & (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) & (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) & (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) | (((e23 = h2(e13)) | (e23 = h2(e12)) | (e23 = h2(e11)) | (e23 = h2(e10))) & ((e22 = h2(e13)) | (e22 = h2(e12)) | (e22 = h2(e11)) | (e22 = h2(e10))) & ((e21 = h2(e13)) | (e21 = h2(e12)) | (e21 = h2(e11)) | (e21 = h2(e10))) & ((e20 = h2(e13)) | (e20 = h2(e12)) | (e20 = h2(e11)) | (e20 = h2(e10))) & (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) & (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) & (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) & (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) & (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) & (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) & (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) & (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) & (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) & (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) & (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) & (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) & (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) & (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) & (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) & (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) | (((e23 = h1(e13)) | (e23 = h1(e12)) | (e23 = h1(e11)) | (e23 = h1(e10))) & ((e22 = h1(e13)) | (e22 = h1(e12)) | (e22 = h1(e11)) | (e22 = h1(e10))) & ((e21 = h1(e13)) | (e21 = h1(e12)) | (e21 = h1(e11)) | (e21 = h1(e10))) & ((e20 = h1(e13)) | (e20 = h1(e12)) | (e20 = h1(e11)) | (e20 = h1(e10))) & (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) & (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) & (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) & (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) & (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) & (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) & (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) & (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) & (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) & (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) & (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) & (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) & (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) & (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) & (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) & (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG109+1.p', co1)).
fof(f2676, plain, (~ spl102_271 | ~ spl102_272), inference(avatar_split_clause, [], [f931, f2673, f2669])).
fof(f931, plain, (~ (e20 = h4(e13)) | ~ sP99), inference(cnf_transformation, [], [f220])).
fof(f220, plain, ((~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10))) | ~ sP99), inference(nnf_transformation, [], [f124])).
fof(f2666, plain, (~ spl102_267 | ~ spl102_270), inference(avatar_split_clause, [], [f925, f2663, f2649])).
fof(f925, plain, (~ (e21 = h4(e11)) | ~ sP100), inference(cnf_transformation, [], [f219])).
fof(f219, plain, ((~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10))) | ~ sP100), inference(nnf_transformation, [], [f125])).
fof(f2641, plain, (~ spl102_263 | ~ spl102_265), inference(avatar_split_clause, [], [f922, f2638, f2629])).
fof(f922, plain, (~ (e22 = h4(e12)) | ~ sP101), inference(cnf_transformation, [], [f218])).
fof(f218, plain, ((~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10))) | ~ sP101), inference(nnf_transformation, [], [f126])).
fof(f2589, plain, spl102_126, inference(avatar_split_clause, [], [f901, f1588])).
fof(f901, plain, (op2(e20, e20) = e21), inference(cnf_transformation, [], [f13])).
fof(f2588, plain, spl102_79, inference(avatar_split_clause, [], [f2587, f1388])).
fof(f2587, plain, (e22 = op2(e23, e20)), inference(forward_demodulation, [], [f902, f903])).
fof(f902, plain, (e22 = op2(op2(e20, op2(e20, e20)), e20)), inference(cnf_transformation, [], [f13])).
fof(f2586, plain, spl102_62, inference(avatar_split_clause, [], [f898, f1284])).
fof(f898, plain, (op1(e10, e10) = e11), inference(cnf_transformation, [], [f12])).
fof(f2585, plain, spl102_15, inference(avatar_split_clause, [], [f2584, f1084])).
fof(f2584, plain, (e12 = op1(e13, e10)), inference(forward_demodulation, [], [f899, f900])).
fof(f899, plain, (e12 = op1(op1(e10, op1(e10, e10)), e10)), inference(cnf_transformation, [], [f12])).
fof(f2583, plain, (spl102_261 | spl102_260 | spl102_259 | spl102_258 | spl102_257 | spl102_256 | spl102_255 | spl102_254 | spl102_253 | spl102_252 | spl102_251 | spl102_250 | spl102_249 | spl102_248 | spl102_247 | spl102_68), inference(avatar_split_clause, [], [f884, f1341, f2418, f2428, f2438, f2448, f2458, f2468, f2478, f2488, f2498, f2508, f2518, f2528, f2538, f2548, f2558])).
fof(f2558, plain, (spl102_261 <=> sP45), introduced(avatar_definition, [new_symbols(naming, [spl102_261])])).
fof(f2548, plain, (spl102_260 <=> sP46), introduced(avatar_definition, [new_symbols(naming, [spl102_260])])).
fof(f2538, plain, (spl102_259 <=> sP47), introduced(avatar_definition, [new_symbols(naming, [spl102_259])])).
fof(f2528, plain, (spl102_258 <=> sP48), introduced(avatar_definition, [new_symbols(naming, [spl102_258])])).
fof(f2518, plain, (spl102_257 <=> sP49), introduced(avatar_definition, [new_symbols(naming, [spl102_257])])).
fof(f2508, plain, (spl102_256 <=> sP50), introduced(avatar_definition, [new_symbols(naming, [spl102_256])])).
fof(f2498, plain, (spl102_255 <=> sP51), introduced(avatar_definition, [new_symbols(naming, [spl102_255])])).
fof(f2488, plain, (spl102_254 <=> sP52), introduced(avatar_definition, [new_symbols(naming, [spl102_254])])).
fof(f2478, plain, (spl102_253 <=> sP53), introduced(avatar_definition, [new_symbols(naming, [spl102_253])])).
fof(f2468, plain, (spl102_252 <=> sP54), introduced(avatar_definition, [new_symbols(naming, [spl102_252])])).
fof(f2458, plain, (spl102_251 <=> sP55), introduced(avatar_definition, [new_symbols(naming, [spl102_251])])).
fof(f2448, plain, (spl102_250 <=> sP56), introduced(avatar_definition, [new_symbols(naming, [spl102_250])])).
fof(f2438, plain, (spl102_249 <=> sP57), introduced(avatar_definition, [new_symbols(naming, [spl102_249])])).
fof(f2428, plain, (spl102_248 <=> sP58), introduced(avatar_definition, [new_symbols(naming, [spl102_248])])).
fof(f2418, plain, (spl102_247 <=> sP59), introduced(avatar_definition, [new_symbols(naming, [spl102_247])])).
fof(f884, plain, ((e23 = op2(e23, e23)) | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45), inference(cnf_transformation, [], [f114])).
fof(f114, plain, ((((((~ (e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23))) & (~ (e22 = op2(e23, e23)) | ~ (e22 = op2(e23, e23))) & (~ (e21 = op2(e23, e23)) | ~ (e21 = op2(e23, e23))) & (~ (e20 = op2(e23, e23)) | ~ (e20 = op2(e23, e23)))) | ~ (e23 = op2(e23, e23))) & (e23 = op2(e23, e23))) | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75) & (((e23 = op2(e23, op2(e23, e23))) & ~ (op2(e23, e23) = op2(e23, e23)) & (op2(e23, e23) = op2(e23, e23))) | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60) & ((~ (e23 = op2(e23, e23)) & ~ (e22 = op2(e22, e23)) & ~ (e21 = op2(e21, e23)) & ~ (e20 = op2(e20, e23)) & (e23 = op2(e23, e23)) & (e23 = op2(e23, e23))) | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45)), inference(definition_folding, [], [f21, e113, e112, e111, e110, e109, e108, e107, e106, e105, e104, e103, e102, e101, e100, e99, e98, e97, e96, e95, e94, e93, e92, e91, e90, e89, e88, e87, e86, e85, e84, e83, e82, e81, e80, e79, e78, e77, e76, e75, e74, e73, e72, e71, e70, e69])).
fof(f69, plain, ((~ (e23 = op2(e23, e20)) & ~ (e22 = op2(e22, e20)) & ~ (e21 = op2(e21, e20)) & ~ (e20 = op2(e20, e20)) & (e20 = op2(e20, e20)) & (e20 = op2(e20, e20))) | ~ sP45), inference(usedef, [], [e69])).
fof(e69, plain, (sP45 <=> (~ (e23 = op2(e23, e20)) & ~ (e22 = op2(e22, e20)) & ~ (e21 = op2(e21, e20)) & ~ (e20 = op2(e20, e20)) & (e20 = op2(e20, e20)) & (e20 = op2(e20, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP45])])).
fof(f70, plain, ((~ (e23 = op2(e23, e20)) & ~ (e22 = op2(e22, e20)) & ~ (e21 = op2(e21, e20)) & ~ (e20 = op2(e20, e20)) & (e20 = op2(e21, e20)) & (e20 = op2(e20, e21))) | ~ sP46), inference(usedef, [], [e70])).
fof(e70, plain, (sP46 <=> (~ (e23 = op2(e23, e20)) & ~ (e22 = op2(e22, e20)) & ~ (e21 = op2(e21, e20)) & ~ (e20 = op2(e20, e20)) & (e20 = op2(e21, e20)) & (e20 = op2(e20, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP46])])).
fof(f71, plain, ((~ (e23 = op2(e23, e20)) & ~ (e22 = op2(e22, e20)) & ~ (e21 = op2(e21, e20)) & ~ (e20 = op2(e20, e20)) & (e20 = op2(e22, e20)) & (e20 = op2(e20, e22))) | ~ sP47), inference(usedef, [], [e71])).
fof(e71, plain, (sP47 <=> (~ (e23 = op2(e23, e20)) & ~ (e22 = op2(e22, e20)) & ~ (e21 = op2(e21, e20)) & ~ (e20 = op2(e20, e20)) & (e20 = op2(e22, e20)) & (e20 = op2(e20, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP47])])).
fof(f72, plain, ((~ (e23 = op2(e23, e20)) & ~ (e22 = op2(e22, e20)) & ~ (e21 = op2(e21, e20)) & ~ (e20 = op2(e20, e20)) & (e20 = op2(e23, e20)) & (e20 = op2(e20, e23))) | ~ sP48), inference(usedef, [], [e72])).
fof(e72, plain, (sP48 <=> (~ (e23 = op2(e23, e20)) & ~ (e22 = op2(e22, e20)) & ~ (e21 = op2(e21, e20)) & ~ (e20 = op2(e20, e20)) & (e20 = op2(e23, e20)) & (e20 = op2(e20, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP48])])).
fof(f73, plain, ((~ (e23 = op2(e23, e21)) & ~ (e22 = op2(e22, e21)) & ~ (e21 = op2(e21, e21)) & ~ (e20 = op2(e20, e21)) & (e21 = op2(e20, e21)) & (e21 = op2(e21, e20))) | ~ sP49), inference(usedef, [], [e73])).
fof(e73, plain, (sP49 <=> (~ (e23 = op2(e23, e21)) & ~ (e22 = op2(e22, e21)) & ~ (e21 = op2(e21, e21)) & ~ (e20 = op2(e20, e21)) & (e21 = op2(e20, e21)) & (e21 = op2(e21, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP49])])).
fof(f74, plain, ((~ (e23 = op2(e23, e21)) & ~ (e22 = op2(e22, e21)) & ~ (e21 = op2(e21, e21)) & ~ (e20 = op2(e20, e21)) & (e21 = op2(e21, e21)) & (e21 = op2(e21, e21))) | ~ sP50), inference(usedef, [], [e74])).
fof(e74, plain, (sP50 <=> (~ (e23 = op2(e23, e21)) & ~ (e22 = op2(e22, e21)) & ~ (e21 = op2(e21, e21)) & ~ (e20 = op2(e20, e21)) & (e21 = op2(e21, e21)) & (e21 = op2(e21, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP50])])).
fof(f75, plain, ((~ (e23 = op2(e23, e21)) & ~ (e22 = op2(e22, e21)) & ~ (e21 = op2(e21, e21)) & ~ (e20 = op2(e20, e21)) & (e21 = op2(e22, e21)) & (e21 = op2(e21, e22))) | ~ sP51), inference(usedef, [], [e75])).
fof(e75, plain, (sP51 <=> (~ (e23 = op2(e23, e21)) & ~ (e22 = op2(e22, e21)) & ~ (e21 = op2(e21, e21)) & ~ (e20 = op2(e20, e21)) & (e21 = op2(e22, e21)) & (e21 = op2(e21, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP51])])).
fof(f76, plain, ((~ (e23 = op2(e23, e21)) & ~ (e22 = op2(e22, e21)) & ~ (e21 = op2(e21, e21)) & ~ (e20 = op2(e20, e21)) & (e21 = op2(e23, e21)) & (e21 = op2(e21, e23))) | ~ sP52), inference(usedef, [], [e76])).
fof(e76, plain, (sP52 <=> (~ (e23 = op2(e23, e21)) & ~ (e22 = op2(e22, e21)) & ~ (e21 = op2(e21, e21)) & ~ (e20 = op2(e20, e21)) & (e21 = op2(e23, e21)) & (e21 = op2(e21, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP52])])).
fof(f77, plain, ((~ (e23 = op2(e23, e22)) & ~ (e22 = op2(e22, e22)) & ~ (e21 = op2(e21, e22)) & ~ (e20 = op2(e20, e22)) & (e22 = op2(e20, e22)) & (e22 = op2(e22, e20))) | ~ sP53), inference(usedef, [], [e77])).
fof(e77, plain, (sP53 <=> (~ (e23 = op2(e23, e22)) & ~ (e22 = op2(e22, e22)) & ~ (e21 = op2(e21, e22)) & ~ (e20 = op2(e20, e22)) & (e22 = op2(e20, e22)) & (e22 = op2(e22, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP53])])).
fof(f78, plain, ((~ (e23 = op2(e23, e22)) & ~ (e22 = op2(e22, e22)) & ~ (e21 = op2(e21, e22)) & ~ (e20 = op2(e20, e22)) & (e22 = op2(e21, e22)) & (e22 = op2(e22, e21))) | ~ sP54), inference(usedef, [], [e78])).
fof(e78, plain, (sP54 <=> (~ (e23 = op2(e23, e22)) & ~ (e22 = op2(e22, e22)) & ~ (e21 = op2(e21, e22)) & ~ (e20 = op2(e20, e22)) & (e22 = op2(e21, e22)) & (e22 = op2(e22, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP54])])).
fof(f79, plain, ((~ (e23 = op2(e23, e22)) & ~ (e22 = op2(e22, e22)) & ~ (e21 = op2(e21, e22)) & ~ (e20 = op2(e20, e22)) & (e22 = op2(e22, e22)) & (e22 = op2(e22, e22))) | ~ sP55), inference(usedef, [], [e79])).
fof(e79, plain, (sP55 <=> (~ (e23 = op2(e23, e22)) & ~ (e22 = op2(e22, e22)) & ~ (e21 = op2(e21, e22)) & ~ (e20 = op2(e20, e22)) & (e22 = op2(e22, e22)) & (e22 = op2(e22, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP55])])).
fof(f80, plain, ((~ (e23 = op2(e23, e22)) & ~ (e22 = op2(e22, e22)) & ~ (e21 = op2(e21, e22)) & ~ (e20 = op2(e20, e22)) & (e22 = op2(e23, e22)) & (e22 = op2(e22, e23))) | ~ sP56), inference(usedef, [], [e80])).
fof(e80, plain, (sP56 <=> (~ (e23 = op2(e23, e22)) & ~ (e22 = op2(e22, e22)) & ~ (e21 = op2(e21, e22)) & ~ (e20 = op2(e20, e22)) & (e22 = op2(e23, e22)) & (e22 = op2(e22, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP56])])).
fof(f81, plain, ((~ (e23 = op2(e23, e23)) & ~ (e22 = op2(e22, e23)) & ~ (e21 = op2(e21, e23)) & ~ (e20 = op2(e20, e23)) & (e23 = op2(e20, e23)) & (e23 = op2(e23, e20))) | ~ sP57), inference(usedef, [], [e81])).
fof(e81, plain, (sP57 <=> (~ (e23 = op2(e23, e23)) & ~ (e22 = op2(e22, e23)) & ~ (e21 = op2(e21, e23)) & ~ (e20 = op2(e20, e23)) & (e23 = op2(e20, e23)) & (e23 = op2(e23, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP57])])).
fof(f82, plain, ((~ (e23 = op2(e23, e23)) & ~ (e22 = op2(e22, e23)) & ~ (e21 = op2(e21, e23)) & ~ (e20 = op2(e20, e23)) & (e23 = op2(e21, e23)) & (e23 = op2(e23, e21))) | ~ sP58), inference(usedef, [], [e82])).
fof(e82, plain, (sP58 <=> (~ (e23 = op2(e23, e23)) & ~ (e22 = op2(e22, e23)) & ~ (e21 = op2(e21, e23)) & ~ (e20 = op2(e20, e23)) & (e23 = op2(e21, e23)) & (e23 = op2(e23, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP58])])).
fof(f83, plain, ((~ (e23 = op2(e23, e23)) & ~ (e22 = op2(e22, e23)) & ~ (e21 = op2(e21, e23)) & ~ (e20 = op2(e20, e23)) & (e23 = op2(e22, e23)) & (e23 = op2(e23, e22))) | ~ sP59), inference(usedef, [], [e83])).
fof(e83, plain, (sP59 <=> (~ (e23 = op2(e23, e23)) & ~ (e22 = op2(e22, e23)) & ~ (e21 = op2(e21, e23)) & ~ (e20 = op2(e20, e23)) & (e23 = op2(e22, e23)) & (e23 = op2(e23, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP59])])).
fof(f84, plain, (((e20 = op2(e20, op2(e20, e20))) & ~ (op2(e20, e20) = op2(e20, e20)) & (op2(e20, e20) = op2(e20, e20))) | ~ sP60), inference(usedef, [], [e84])).
fof(e84, plain, (sP60 <=> ((e20 = op2(e20, op2(e20, e20))) & ~ (op2(e20, e20) = op2(e20, e20)) & (op2(e20, e20) = op2(e20, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP60])])).
fof(f85, plain, (((e21 = op2(e21, op2(e20, e21))) & ~ (op2(e20, e21) = op2(e21, e21)) & (op2(e20, e21) = op2(e21, e20))) | ~ sP61), inference(usedef, [], [e85])).
fof(e85, plain, (sP61 <=> ((e21 = op2(e21, op2(e20, e21))) & ~ (op2(e20, e21) = op2(e21, e21)) & (op2(e20, e21) = op2(e21, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP61])])).
fof(f86, plain, (((e22 = op2(e22, op2(e20, e22))) & ~ (op2(e20, e22) = op2(e22, e22)) & (op2(e20, e22) = op2(e22, e20))) | ~ sP62), inference(usedef, [], [e86])).
fof(e86, plain, (sP62 <=> ((e22 = op2(e22, op2(e20, e22))) & ~ (op2(e20, e22) = op2(e22, e22)) & (op2(e20, e22) = op2(e22, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP62])])).
fof(f87, plain, (((e23 = op2(e23, op2(e20, e23))) & ~ (op2(e20, e23) = op2(e23, e23)) & (op2(e20, e23) = op2(e23, e20))) | ~ sP63), inference(usedef, [], [e87])).
fof(e87, plain, (sP63 <=> ((e23 = op2(e23, op2(e20, e23))) & ~ (op2(e20, e23) = op2(e23, e23)) & (op2(e20, e23) = op2(e23, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP63])])).
fof(f88, plain, (((e20 = op2(e20, op2(e21, e20))) & ~ (op2(e20, e20) = op2(e21, e20)) & (op2(e20, e21) = op2(e21, e20))) | ~ sP64), inference(usedef, [], [e88])).
fof(e88, plain, (sP64 <=> ((e20 = op2(e20, op2(e21, e20))) & ~ (op2(e20, e20) = op2(e21, e20)) & (op2(e20, e21) = op2(e21, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP64])])).
fof(f89, plain, (((e21 = op2(e21, op2(e21, e21))) & ~ (op2(e21, e21) = op2(e21, e21)) & (op2(e21, e21) = op2(e21, e21))) | ~ sP65), inference(usedef, [], [e89])).
fof(e89, plain, (sP65 <=> ((e21 = op2(e21, op2(e21, e21))) & ~ (op2(e21, e21) = op2(e21, e21)) & (op2(e21, e21) = op2(e21, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP65])])).
fof(f90, plain, (((e22 = op2(e22, op2(e21, e22))) & ~ (op2(e21, e22) = op2(e22, e22)) & (op2(e21, e22) = op2(e22, e21))) | ~ sP66), inference(usedef, [], [e90])).
fof(e90, plain, (sP66 <=> ((e22 = op2(e22, op2(e21, e22))) & ~ (op2(e21, e22) = op2(e22, e22)) & (op2(e21, e22) = op2(e22, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP66])])).
fof(f91, plain, (((e23 = op2(e23, op2(e21, e23))) & ~ (op2(e21, e23) = op2(e23, e23)) & (op2(e21, e23) = op2(e23, e21))) | ~ sP67), inference(usedef, [], [e91])).
fof(e91, plain, (sP67 <=> ((e23 = op2(e23, op2(e21, e23))) & ~ (op2(e21, e23) = op2(e23, e23)) & (op2(e21, e23) = op2(e23, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP67])])).
fof(f92, plain, (((e20 = op2(e20, op2(e22, e20))) & ~ (op2(e20, e20) = op2(e22, e20)) & (op2(e20, e22) = op2(e22, e20))) | ~ sP68), inference(usedef, [], [e92])).
fof(e92, plain, (sP68 <=> ((e20 = op2(e20, op2(e22, e20))) & ~ (op2(e20, e20) = op2(e22, e20)) & (op2(e20, e22) = op2(e22, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP68])])).
fof(f93, plain, (((e21 = op2(e21, op2(e22, e21))) & ~ (op2(e21, e21) = op2(e22, e21)) & (op2(e21, e22) = op2(e22, e21))) | ~ sP69), inference(usedef, [], [e93])).
fof(e93, plain, (sP69 <=> ((e21 = op2(e21, op2(e22, e21))) & ~ (op2(e21, e21) = op2(e22, e21)) & (op2(e21, e22) = op2(e22, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP69])])).
fof(f94, plain, (((e22 = op2(e22, op2(e22, e22))) & ~ (op2(e22, e22) = op2(e22, e22)) & (op2(e22, e22) = op2(e22, e22))) | ~ sP70), inference(usedef, [], [e94])).
fof(e94, plain, (sP70 <=> ((e22 = op2(e22, op2(e22, e22))) & ~ (op2(e22, e22) = op2(e22, e22)) & (op2(e22, e22) = op2(e22, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP70])])).
fof(f95, plain, (((e23 = op2(e23, op2(e22, e23))) & ~ (op2(e22, e23) = op2(e23, e23)) & (op2(e22, e23) = op2(e23, e22))) | ~ sP71), inference(usedef, [], [e95])).
fof(e95, plain, (sP71 <=> ((e23 = op2(e23, op2(e22, e23))) & ~ (op2(e22, e23) = op2(e23, e23)) & (op2(e22, e23) = op2(e23, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP71])])).
fof(f96, plain, (((e20 = op2(e20, op2(e23, e20))) & ~ (op2(e20, e20) = op2(e23, e20)) & (op2(e20, e23) = op2(e23, e20))) | ~ sP72), inference(usedef, [], [e96])).
fof(e96, plain, (sP72 <=> ((e20 = op2(e20, op2(e23, e20))) & ~ (op2(e20, e20) = op2(e23, e20)) & (op2(e20, e23) = op2(e23, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP72])])).
fof(f97, plain, (((e21 = op2(e21, op2(e23, e21))) & ~ (op2(e21, e21) = op2(e23, e21)) & (op2(e21, e23) = op2(e23, e21))) | ~ sP73), inference(usedef, [], [e97])).
fof(e97, plain, (sP73 <=> ((e21 = op2(e21, op2(e23, e21))) & ~ (op2(e21, e21) = op2(e23, e21)) & (op2(e21, e23) = op2(e23, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP73])])).
fof(f98, plain, (((e22 = op2(e22, op2(e23, e22))) & ~ (op2(e22, e22) = op2(e23, e22)) & (op2(e22, e23) = op2(e23, e22))) | ~ sP74), inference(usedef, [], [e98])).
fof(e98, plain, (sP74 <=> ((e22 = op2(e22, op2(e23, e22))) & ~ (op2(e22, e22) = op2(e23, e22)) & (op2(e22, e23) = op2(e23, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP74])])).
fof(f99, plain, (((((~ (op2(e20, e20) = e23) | ~ (op2(e20, e20) = e23)) & (~ (op2(e20, e20) = e22) | ~ (op2(e20, e20) = e22)) & (~ (op2(e20, e20) = e21) | ~ (op2(e20, e20) = e21)) & (~ (e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20)))) | ~ (e20 = op2(e20, e20))) & (e20 = op2(e20, e20))) | ~ sP75), inference(usedef, [], [e99])).
fof(e99, plain, (sP75 <=> ((((~ (op2(e20, e20) = e23) | ~ (op2(e20, e20) = e23)) & (~ (op2(e20, e20) = e22) | ~ (op2(e20, e20) = e22)) & (~ (op2(e20, e20) = e21) | ~ (op2(e20, e20) = e21)) & (~ (e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20)))) | ~ (e20 = op2(e20, e20))) & (e20 = op2(e20, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP75])])).
fof(f100, plain, (((((~ (e23 = op2(e21, e20)) | ~ (e23 = op2(e20, e21))) & (~ (e22 = op2(e21, e20)) | ~ (e22 = op2(e20, e21))) & (~ (e21 = op2(e21, e20)) | ~ (e21 = op2(e20, e21))) & (~ (e20 = op2(e21, e20)) | ~ (e20 = op2(e20, e21)))) | ~ (op2(e20, e20) = e21)) & (op2(e20, e20) = e21)) | ~ sP76), inference(usedef, [], [e100])).
fof(e100, plain, (sP76 <=> ((((~ (e23 = op2(e21, e20)) | ~ (e23 = op2(e20, e21))) & (~ (e22 = op2(e21, e20)) | ~ (e22 = op2(e20, e21))) & (~ (e21 = op2(e21, e20)) | ~ (e21 = op2(e20, e21))) & (~ (e20 = op2(e21, e20)) | ~ (e20 = op2(e20, e21)))) | ~ (op2(e20, e20) = e21)) & (op2(e20, e20) = e21))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP76])])).
fof(f101, plain, (((((~ (e23 = op2(e22, e20)) | ~ (e23 = op2(e20, e22))) & (~ (e22 = op2(e22, e20)) | ~ (e22 = op2(e20, e22))) & (~ (e21 = op2(e22, e20)) | ~ (e21 = op2(e20, e22))) & (~ (e20 = op2(e22, e20)) | ~ (e20 = op2(e20, e22)))) | ~ (op2(e20, e20) = e22)) & (op2(e20, e20) = e22)) | ~ sP77), inference(usedef, [], [e101])).
fof(e101, plain, (sP77 <=> ((((~ (e23 = op2(e22, e20)) | ~ (e23 = op2(e20, e22))) & (~ (e22 = op2(e22, e20)) | ~ (e22 = op2(e20, e22))) & (~ (e21 = op2(e22, e20)) | ~ (e21 = op2(e20, e22))) & (~ (e20 = op2(e22, e20)) | ~ (e20 = op2(e20, e22)))) | ~ (op2(e20, e20) = e22)) & (op2(e20, e20) = e22))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP77])])).
fof(f102, plain, (((((~ (e23 = op2(e23, e20)) | ~ (e23 = op2(e20, e23))) & (~ (e22 = op2(e23, e20)) | ~ (e22 = op2(e20, e23))) & (~ (e21 = op2(e23, e20)) | ~ (e21 = op2(e20, e23))) & (~ (e20 = op2(e23, e20)) | ~ (e20 = op2(e20, e23)))) | ~ (op2(e20, e20) = e23)) & (op2(e20, e20) = e23)) | ~ sP78), inference(usedef, [], [e102])).
fof(e102, plain, (sP78 <=> ((((~ (e23 = op2(e23, e20)) | ~ (e23 = op2(e20, e23))) & (~ (e22 = op2(e23, e20)) | ~ (e22 = op2(e20, e23))) & (~ (e21 = op2(e23, e20)) | ~ (e21 = op2(e20, e23))) & (~ (e20 = op2(e23, e20)) | ~ (e20 = op2(e20, e23)))) | ~ (op2(e20, e20) = e23)) & (op2(e20, e20) = e23))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP78])])).
fof(f103, plain, (((((~ (e23 = op2(e20, e21)) | ~ (e23 = op2(e21, e20))) & (~ (e22 = op2(e20, e21)) | ~ (e22 = op2(e21, e20))) & (~ (e21 = op2(e20, e21)) | ~ (e21 = op2(e21, e20))) & (~ (e20 = op2(e20, e21)) | ~ (e20 = op2(e21, e20)))) | ~ (e20 = op2(e21, e21))) & (e20 = op2(e21, e21))) | ~ sP79), inference(usedef, [], [e103])).
fof(e103, plain, (sP79 <=> ((((~ (e23 = op2(e20, e21)) | ~ (e23 = op2(e21, e20))) & (~ (e22 = op2(e20, e21)) | ~ (e22 = op2(e21, e20))) & (~ (e21 = op2(e20, e21)) | ~ (e21 = op2(e21, e20))) & (~ (e20 = op2(e20, e21)) | ~ (e20 = op2(e21, e20)))) | ~ (e20 = op2(e21, e21))) & (e20 = op2(e21, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP79])])).
fof(f104, plain, (((((~ (e23 = op2(e21, e21)) | ~ (e23 = op2(e21, e21))) & (~ (e22 = op2(e21, e21)) | ~ (e22 = op2(e21, e21))) & (~ (e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & (~ (e20 = op2(e21, e21)) | ~ (e20 = op2(e21, e21)))) | ~ (e21 = op2(e21, e21))) & (e21 = op2(e21, e21))) | ~ sP80), inference(usedef, [], [e104])).
fof(e104, plain, (sP80 <=> ((((~ (e23 = op2(e21, e21)) | ~ (e23 = op2(e21, e21))) & (~ (e22 = op2(e21, e21)) | ~ (e22 = op2(e21, e21))) & (~ (e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & (~ (e20 = op2(e21, e21)) | ~ (e20 = op2(e21, e21)))) | ~ (e21 = op2(e21, e21))) & (e21 = op2(e21, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP80])])).
fof(f105, plain, (((((~ (e23 = op2(e22, e21)) | ~ (e23 = op2(e21, e22))) & (~ (e22 = op2(e22, e21)) | ~ (e22 = op2(e21, e22))) & (~ (e21 = op2(e22, e21)) | ~ (e21 = op2(e21, e22))) & (~ (e20 = op2(e22, e21)) | ~ (e20 = op2(e21, e22)))) | ~ (e22 = op2(e21, e21))) & (e22 = op2(e21, e21))) | ~ sP81), inference(usedef, [], [e105])).
fof(e105, plain, (sP81 <=> ((((~ (e23 = op2(e22, e21)) | ~ (e23 = op2(e21, e22))) & (~ (e22 = op2(e22, e21)) | ~ (e22 = op2(e21, e22))) & (~ (e21 = op2(e22, e21)) | ~ (e21 = op2(e21, e22))) & (~ (e20 = op2(e22, e21)) | ~ (e20 = op2(e21, e22)))) | ~ (e22 = op2(e21, e21))) & (e22 = op2(e21, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP81])])).
fof(f106, plain, (((((~ (e23 = op2(e23, e21)) | ~ (e23 = op2(e21, e23))) & (~ (e22 = op2(e23, e21)) | ~ (e22 = op2(e21, e23))) & (~ (e21 = op2(e23, e21)) | ~ (e21 = op2(e21, e23))) & (~ (e20 = op2(e23, e21)) | ~ (e20 = op2(e21, e23)))) | ~ (e23 = op2(e21, e21))) & (e23 = op2(e21, e21))) | ~ sP82), inference(usedef, [], [e106])).
fof(e106, plain, (sP82 <=> ((((~ (e23 = op2(e23, e21)) | ~ (e23 = op2(e21, e23))) & (~ (e22 = op2(e23, e21)) | ~ (e22 = op2(e21, e23))) & (~ (e21 = op2(e23, e21)) | ~ (e21 = op2(e21, e23))) & (~ (e20 = op2(e23, e21)) | ~ (e20 = op2(e21, e23)))) | ~ (e23 = op2(e21, e21))) & (e23 = op2(e21, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP82])])).
fof(f107, plain, (((((~ (e23 = op2(e20, e22)) | ~ (e23 = op2(e22, e20))) & (~ (e22 = op2(e20, e22)) | ~ (e22 = op2(e22, e20))) & (~ (e21 = op2(e20, e22)) | ~ (e21 = op2(e22, e20))) & (~ (e20 = op2(e20, e22)) | ~ (e20 = op2(e22, e20)))) | ~ (e20 = op2(e22, e22))) & (e20 = op2(e22, e22))) | ~ sP83), inference(usedef, [], [e107])).
fof(e107, plain, (sP83 <=> ((((~ (e23 = op2(e20, e22)) | ~ (e23 = op2(e22, e20))) & (~ (e22 = op2(e20, e22)) | ~ (e22 = op2(e22, e20))) & (~ (e21 = op2(e20, e22)) | ~ (e21 = op2(e22, e20))) & (~ (e20 = op2(e20, e22)) | ~ (e20 = op2(e22, e20)))) | ~ (e20 = op2(e22, e22))) & (e20 = op2(e22, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP83])])).
fof(f108, plain, (((((~ (e23 = op2(e21, e22)) | ~ (e23 = op2(e22, e21))) & (~ (e22 = op2(e21, e22)) | ~ (e22 = op2(e22, e21))) & (~ (e21 = op2(e21, e22)) | ~ (e21 = op2(e22, e21))) & (~ (e20 = op2(e21, e22)) | ~ (e20 = op2(e22, e21)))) | ~ (e21 = op2(e22, e22))) & (e21 = op2(e22, e22))) | ~ sP84), inference(usedef, [], [e108])).
fof(e108, plain, (sP84 <=> ((((~ (e23 = op2(e21, e22)) | ~ (e23 = op2(e22, e21))) & (~ (e22 = op2(e21, e22)) | ~ (e22 = op2(e22, e21))) & (~ (e21 = op2(e21, e22)) | ~ (e21 = op2(e22, e21))) & (~ (e20 = op2(e21, e22)) | ~ (e20 = op2(e22, e21)))) | ~ (e21 = op2(e22, e22))) & (e21 = op2(e22, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP84])])).
fof(f109, plain, (((((~ (e23 = op2(e22, e22)) | ~ (e23 = op2(e22, e22))) & (~ (e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & (~ (e21 = op2(e22, e22)) | ~ (e21 = op2(e22, e22))) & (~ (e20 = op2(e22, e22)) | ~ (e20 = op2(e22, e22)))) | ~ (e22 = op2(e22, e22))) & (e22 = op2(e22, e22))) | ~ sP85), inference(usedef, [], [e109])).
fof(e109, plain, (sP85 <=> ((((~ (e23 = op2(e22, e22)) | ~ (e23 = op2(e22, e22))) & (~ (e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & (~ (e21 = op2(e22, e22)) | ~ (e21 = op2(e22, e22))) & (~ (e20 = op2(e22, e22)) | ~ (e20 = op2(e22, e22)))) | ~ (e22 = op2(e22, e22))) & (e22 = op2(e22, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP85])])).
fof(f110, plain, (((((~ (e23 = op2(e23, e22)) | ~ (e23 = op2(e22, e23))) & (~ (e22 = op2(e23, e22)) | ~ (e22 = op2(e22, e23))) & (~ (e21 = op2(e23, e22)) | ~ (e21 = op2(e22, e23))) & (~ (e20 = op2(e23, e22)) | ~ (e20 = op2(e22, e23)))) | ~ (e23 = op2(e22, e22))) & (e23 = op2(e22, e22))) | ~ sP86), inference(usedef, [], [e110])).
fof(e110, plain, (sP86 <=> ((((~ (e23 = op2(e23, e22)) | ~ (e23 = op2(e22, e23))) & (~ (e22 = op2(e23, e22)) | ~ (e22 = op2(e22, e23))) & (~ (e21 = op2(e23, e22)) | ~ (e21 = op2(e22, e23))) & (~ (e20 = op2(e23, e22)) | ~ (e20 = op2(e22, e23)))) | ~ (e23 = op2(e22, e22))) & (e23 = op2(e22, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP86])])).
fof(f111, plain, (((((~ (e23 = op2(e20, e23)) | ~ (e23 = op2(e23, e20))) & (~ (e22 = op2(e20, e23)) | ~ (e22 = op2(e23, e20))) & (~ (e21 = op2(e20, e23)) | ~ (e21 = op2(e23, e20))) & (~ (e20 = op2(e20, e23)) | ~ (e20 = op2(e23, e20)))) | ~ (e20 = op2(e23, e23))) & (e20 = op2(e23, e23))) | ~ sP87), inference(usedef, [], [e111])).
fof(e111, plain, (sP87 <=> ((((~ (e23 = op2(e20, e23)) | ~ (e23 = op2(e23, e20))) & (~ (e22 = op2(e20, e23)) | ~ (e22 = op2(e23, e20))) & (~ (e21 = op2(e20, e23)) | ~ (e21 = op2(e23, e20))) & (~ (e20 = op2(e20, e23)) | ~ (e20 = op2(e23, e20)))) | ~ (e20 = op2(e23, e23))) & (e20 = op2(e23, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP87])])).
fof(f112, plain, (((((~ (e23 = op2(e21, e23)) | ~ (e23 = op2(e23, e21))) & (~ (e22 = op2(e21, e23)) | ~ (e22 = op2(e23, e21))) & (~ (e21 = op2(e21, e23)) | ~ (e21 = op2(e23, e21))) & (~ (e20 = op2(e21, e23)) | ~ (e20 = op2(e23, e21)))) | ~ (e21 = op2(e23, e23))) & (e21 = op2(e23, e23))) | ~ sP88), inference(usedef, [], [e112])).
fof(e112, plain, (sP88 <=> ((((~ (e23 = op2(e21, e23)) | ~ (e23 = op2(e23, e21))) & (~ (e22 = op2(e21, e23)) | ~ (e22 = op2(e23, e21))) & (~ (e21 = op2(e21, e23)) | ~ (e21 = op2(e23, e21))) & (~ (e20 = op2(e21, e23)) | ~ (e20 = op2(e23, e21)))) | ~ (e21 = op2(e23, e23))) & (e21 = op2(e23, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP88])])).
fof(f113, plain, (((((~ (e23 = op2(e22, e23)) | ~ (e23 = op2(e23, e22))) & (~ (e22 = op2(e22, e23)) | ~ (e22 = op2(e23, e22))) & (~ (e21 = op2(e22, e23)) | ~ (e21 = op2(e23, e22))) & (~ (e20 = op2(e22, e23)) | ~ (e20 = op2(e23, e22)))) | ~ (e22 = op2(e23, e23))) & (e22 = op2(e23, e23))) | ~ sP89), inference(usedef, [], [e113])).
fof(e113, plain, (sP89 <=> ((((~ (e23 = op2(e22, e23)) | ~ (e23 = op2(e23, e22))) & (~ (e22 = op2(e22, e23)) | ~ (e22 = op2(e23, e22))) & (~ (e21 = op2(e22, e23)) | ~ (e21 = op2(e23, e22))) & (~ (e20 = op2(e22, e23)) | ~ (e20 = op2(e23, e22)))) | ~ (e22 = op2(e23, e23))) & (e22 = op2(e23, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP89])])).
fof(f21, plain, ((((((~ (e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23))) & (~ (e22 = op2(e23, e23)) | ~ (e22 = op2(e23, e23))) & (~ (e21 = op2(e23, e23)) | ~ (e21 = op2(e23, e23))) & (~ (e20 = op2(e23, e23)) | ~ (e20 = op2(e23, e23)))) | ~ (e23 = op2(e23, e23))) & (e23 = op2(e23, e23))) | ((((~ (e23 = op2(e22, e23)) | ~ (e23 = op2(e23, e22))) & (~ (e22 = op2(e22, e23)) | ~ (e22 = op2(e23, e22))) & (~ (e21 = op2(e22, e23)) | ~ (e21 = op2(e23, e22))) & (~ (e20 = op2(e22, e23)) | ~ (e20 = op2(e23, e22)))) | ~ (e22 = op2(e23, e23))) & (e22 = op2(e23, e23))) | ((((~ (e23 = op2(e21, e23)) | ~ (e23 = op2(e23, e21))) & (~ (e22 = op2(e21, e23)) | ~ (e22 = op2(e23, e21))) & (~ (e21 = op2(e21, e23)) | ~ (e21 = op2(e23, e21))) & (~ (e20 = op2(e21, e23)) | ~ (e20 = op2(e23, e21)))) | ~ (e21 = op2(e23, e23))) & (e21 = op2(e23, e23))) | ((((~ (e23 = op2(e20, e23)) | ~ (e23 = op2(e23, e20))) & (~ (e22 = op2(e20, e23)) | ~ (e22 = op2(e23, e20))) & (~ (e21 = op2(e20, e23)) | ~ (e21 = op2(e23, e20))) & (~ (e20 = op2(e20, e23)) | ~ (e20 = op2(e23, e20)))) | ~ (e20 = op2(e23, e23))) & (e20 = op2(e23, e23))) | ((((~ (e23 = op2(e23, e22)) | ~ (e23 = op2(e22, e23))) & (~ (e22 = op2(e23, e22)) | ~ (e22 = op2(e22, e23))) & (~ (e21 = op2(e23, e22)) | ~ (e21 = op2(e22, e23))) & (~ (e20 = op2(e23, e22)) | ~ (e20 = op2(e22, e23)))) | ~ (e23 = op2(e22, e22))) & (e23 = op2(e22, e22))) | ((((~ (e23 = op2(e22, e22)) | ~ (e23 = op2(e22, e22))) & (~ (e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & (~ (e21 = op2(e22, e22)) | ~ (e21 = op2(e22, e22))) & (~ (e20 = op2(e22, e22)) | ~ (e20 = op2(e22, e22)))) | ~ (e22 = op2(e22, e22))) & (e22 = op2(e22, e22))) | ((((~ (e23 = op2(e21, e22)) | ~ (e23 = op2(e22, e21))) & (~ (e22 = op2(e21, e22)) | ~ (e22 = op2(e22, e21))) & (~ (e21 = op2(e21, e22)) | ~ (e21 = op2(e22, e21))) & (~ (e20 = op2(e21, e22)) | ~ (e20 = op2(e22, e21)))) | ~ (e21 = op2(e22, e22))) & (e21 = op2(e22, e22))) | ((((~ (e23 = op2(e20, e22)) | ~ (e23 = op2(e22, e20))) & (~ (e22 = op2(e20, e22)) | ~ (e22 = op2(e22, e20))) & (~ (e21 = op2(e20, e22)) | ~ (e21 = op2(e22, e20))) & (~ (e20 = op2(e20, e22)) | ~ (e20 = op2(e22, e20)))) | ~ (e20 = op2(e22, e22))) & (e20 = op2(e22, e22))) | ((((~ (e23 = op2(e23, e21)) | ~ (e23 = op2(e21, e23))) & (~ (e22 = op2(e23, e21)) | ~ (e22 = op2(e21, e23))) & (~ (e21 = op2(e23, e21)) | ~ (e21 = op2(e21, e23))) & (~ (e20 = op2(e23, e21)) | ~ (e20 = op2(e21, e23)))) | ~ (e23 = op2(e21, e21))) & (e23 = op2(e21, e21))) | ((((~ (e23 = op2(e22, e21)) | ~ (e23 = op2(e21, e22))) & (~ (e22 = op2(e22, e21)) | ~ (e22 = op2(e21, e22))) & (~ (e21 = op2(e22, e21)) | ~ (e21 = op2(e21, e22))) & (~ (e20 = op2(e22, e21)) | ~ (e20 = op2(e21, e22)))) | ~ (e22 = op2(e21, e21))) & (e22 = op2(e21, e21))) | ((((~ (e23 = op2(e21, e21)) | ~ (e23 = op2(e21, e21))) & (~ (e22 = op2(e21, e21)) | ~ (e22 = op2(e21, e21))) & (~ (e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & (~ (e20 = op2(e21, e21)) | ~ (e20 = op2(e21, e21)))) | ~ (e21 = op2(e21, e21))) & (e21 = op2(e21, e21))) | ((((~ (e23 = op2(e20, e21)) | ~ (e23 = op2(e21, e20))) & (~ (e22 = op2(e20, e21)) | ~ (e22 = op2(e21, e20))) & (~ (e21 = op2(e20, e21)) | ~ (e21 = op2(e21, e20))) & (~ (e20 = op2(e20, e21)) | ~ (e20 = op2(e21, e20)))) | ~ (e20 = op2(e21, e21))) & (e20 = op2(e21, e21))) | ((((~ (e23 = op2(e23, e20)) | ~ (e23 = op2(e20, e23))) & (~ (e22 = op2(e23, e20)) | ~ (e22 = op2(e20, e23))) & (~ (e21 = op2(e23, e20)) | ~ (e21 = op2(e20, e23))) & (~ (e20 = op2(e23, e20)) | ~ (e20 = op2(e20, e23)))) | ~ (op2(e20, e20) = e23)) & (op2(e20, e20) = e23)) | ((((~ (e23 = op2(e22, e20)) | ~ (e23 = op2(e20, e22))) & (~ (e22 = op2(e22, e20)) | ~ (e22 = op2(e20, e22))) & (~ (e21 = op2(e22, e20)) | ~ (e21 = op2(e20, e22))) & (~ (e20 = op2(e22, e20)) | ~ (e20 = op2(e20, e22)))) | ~ (op2(e20, e20) = e22)) & (op2(e20, e20) = e22)) | ((((~ (e23 = op2(e21, e20)) | ~ (e23 = op2(e20, e21))) & (~ (e22 = op2(e21, e20)) | ~ (e22 = op2(e20, e21))) & (~ (e21 = op2(e21, e20)) | ~ (e21 = op2(e20, e21))) & (~ (e20 = op2(e21, e20)) | ~ (e20 = op2(e20, e21)))) | ~ (op2(e20, e20) = e21)) & (op2(e20, e20) = e21)) | ((((~ (op2(e20, e20) = e23) | ~ (op2(e20, e20) = e23)) & (~ (op2(e20, e20) = e22) | ~ (op2(e20, e20) = e22)) & (~ (op2(e20, e20) = e21) | ~ (op2(e20, e20) = e21)) & (~ (e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20)))) | ~ (e20 = op2(e20, e20))) & (e20 = op2(e20, e20)))) & (((e23 = op2(e23, op2(e23, e23))) & ~ (op2(e23, e23) = op2(e23, e23)) & (op2(e23, e23) = op2(e23, e23))) | ((e22 = op2(e22, op2(e23, e22))) & ~ (op2(e22, e22) = op2(e23, e22)) & (op2(e22, e23) = op2(e23, e22))) | ((e21 = op2(e21, op2(e23, e21))) & ~ (op2(e21, e21) = op2(e23, e21)) & (op2(e21, e23) = op2(e23, e21))) | ((e20 = op2(e20, op2(e23, e20))) & ~ (op2(e20, e20) = op2(e23, e20)) & (op2(e20, e23) = op2(e23, e20))) | ((e23 = op2(e23, op2(e22, e23))) & ~ (op2(e22, e23) = op2(e23, e23)) & (op2(e22, e23) = op2(e23, e22))) | ((e22 = op2(e22, op2(e22, e22))) & ~ (op2(e22, e22) = op2(e22, e22)) & (op2(e22, e22) = op2(e22, e22))) | ((e21 = op2(e21, op2(e22, e21))) & ~ (op2(e21, e21) = op2(e22, e21)) & (op2(e21, e22) = op2(e22, e21))) | ((e20 = op2(e20, op2(e22, e20))) & ~ (op2(e20, e20) = op2(e22, e20)) & (op2(e20, e22) = op2(e22, e20))) | ((e23 = op2(e23, op2(e21, e23))) & ~ (op2(e21, e23) = op2(e23, e23)) & (op2(e21, e23) = op2(e23, e21))) | ((e22 = op2(e22, op2(e21, e22))) & ~ (op2(e21, e22) = op2(e22, e22)) & (op2(e21, e22) = op2(e22, e21))) | ((e21 = op2(e21, op2(e21, e21))) & ~ (op2(e21, e21) = op2(e21, e21)) & (op2(e21, e21) = op2(e21, e21))) | ((e20 = op2(e20, op2(e21, e20))) & ~ (op2(e20, e20) = op2(e21, e20)) & (op2(e20, e21) = op2(e21, e20))) | ((e23 = op2(e23, op2(e20, e23))) & ~ (op2(e20, e23) = op2(e23, e23)) & (op2(e20, e23) = op2(e23, e20))) | ((e22 = op2(e22, op2(e20, e22))) & ~ (op2(e20, e22) = op2(e22, e22)) & (op2(e20, e22) = op2(e22, e20))) | ((e21 = op2(e21, op2(e20, e21))) & ~ (op2(e20, e21) = op2(e21, e21)) & (op2(e20, e21) = op2(e21, e20))) | ((e20 = op2(e20, op2(e20, e20))) & ~ (op2(e20, e20) = op2(e20, e20)) & (op2(e20, e20) = op2(e20, e20)))) & ((~ (e23 = op2(e23, e23)) & ~ (e22 = op2(e22, e23)) & ~ (e21 = op2(e21, e23)) & ~ (e20 = op2(e20, e23)) & (e23 = op2(e23, e23)) & (e23 = op2(e23, e23))) | (~ (e23 = op2(e23, e23)) & ~ (e22 = op2(e22, e23)) & ~ (e21 = op2(e21, e23)) & ~ (e20 = op2(e20, e23)) & (e23 = op2(e22, e23)) & (e23 = op2(e23, e22))) | (~ (e23 = op2(e23, e23)) & ~ (e22 = op2(e22, e23)) & ~ (e21 = op2(e21, e23)) & ~ (e20 = op2(e20, e23)) & (e23 = op2(e21, e23)) & (e23 = op2(e23, e21))) | (~ (e23 = op2(e23, e23)) & ~ (e22 = op2(e22, e23)) & ~ (e21 = op2(e21, e23)) & ~ (e20 = op2(e20, e23)) & (e23 = op2(e20, e23)) & (e23 = op2(e23, e20))) | (~ (e23 = op2(e23, e22)) & ~ (e22 = op2(e22, e22)) & ~ (e21 = op2(e21, e22)) & ~ (e20 = op2(e20, e22)) & (e22 = op2(e23, e22)) & (e22 = op2(e22, e23))) | (~ (e23 = op2(e23, e22)) & ~ (e22 = op2(e22, e22)) & ~ (e21 = op2(e21, e22)) & ~ (e20 = op2(e20, e22)) & (e22 = op2(e22, e22)) & (e22 = op2(e22, e22))) | (~ (e23 = op2(e23, e22)) & ~ (e22 = op2(e22, e22)) & ~ (e21 = op2(e21, e22)) & ~ (e20 = op2(e20, e22)) & (e22 = op2(e21, e22)) & (e22 = op2(e22, e21))) | (~ (e23 = op2(e23, e22)) & ~ (e22 = op2(e22, e22)) & ~ (e21 = op2(e21, e22)) & ~ (e20 = op2(e20, e22)) & (e22 = op2(e20, e22)) & (e22 = op2(e22, e20))) | (~ (e23 = op2(e23, e21)) & ~ (e22 = op2(e22, e21)) & ~ (e21 = op2(e21, e21)) & ~ (e20 = op2(e20, e21)) & (e21 = op2(e23, e21)) & (e21 = op2(e21, e23))) | (~ (e23 = op2(e23, e21)) & ~ (e22 = op2(e22, e21)) & ~ (e21 = op2(e21, e21)) & ~ (e20 = op2(e20, e21)) & (e21 = op2(e22, e21)) & (e21 = op2(e21, e22))) | (~ (e23 = op2(e23, e21)) & ~ (e22 = op2(e22, e21)) & ~ (e21 = op2(e21, e21)) & ~ (e20 = op2(e20, e21)) & (e21 = op2(e21, e21)) & (e21 = op2(e21, e21))) | (~ (e23 = op2(e23, e21)) & ~ (e22 = op2(e22, e21)) & ~ (e21 = op2(e21, e21)) & ~ (e20 = op2(e20, e21)) & (e21 = op2(e20, e21)) & (e21 = op2(e21, e20))) | (~ (e23 = op2(e23, e20)) & ~ (e22 = op2(e22, e20)) & ~ (e21 = op2(e21, e20)) & ~ (e20 = op2(e20, e20)) & (e20 = op2(e23, e20)) & (e20 = op2(e20, e23))) | (~ (e23 = op2(e23, e20)) & ~ (e22 = op2(e22, e20)) & ~ (e21 = op2(e21, e20)) & ~ (e20 = op2(e20, e20)) & (e20 = op2(e22, e20)) & (e20 = op2(e20, e22))) | (~ (e23 = op2(e23, e20)) & ~ (e22 = op2(e22, e20)) & ~ (e21 = op2(e21, e20)) & ~ (e20 = op2(e20, e20)) & (e20 = op2(e21, e20)) & (e20 = op2(e20, e21))) | (~ (e23 = op2(e23, e20)) & ~ (e22 = op2(e22, e20)) & ~ (e21 = op2(e21, e20)) & ~ (e20 = op2(e20, e20)) & (e20 = op2(e20, e20)) & (e20 = op2(e20, e20))))), inference(ennf_transformation, [], [f11])).
fof(f11, plain, (((~ ((((e23 = op2(e23, e23)) & (e23 = op2(e23, e23))) | ((e22 = op2(e23, e23)) & (e22 = op2(e23, e23))) | ((e21 = op2(e23, e23)) & (e21 = op2(e23, e23))) | ((e20 = op2(e23, e23)) & (e20 = op2(e23, e23)))) & (e23 = op2(e23, e23))) & (e23 = op2(e23, e23))) | (~ ((((e23 = op2(e22, e23)) & (e23 = op2(e23, e22))) | ((e22 = op2(e22, e23)) & (e22 = op2(e23, e22))) | ((e21 = op2(e22, e23)) & (e21 = op2(e23, e22))) | ((e20 = op2(e22, e23)) & (e20 = op2(e23, e22)))) & (e22 = op2(e23, e23))) & (e22 = op2(e23, e23))) | (~ ((((e23 = op2(e21, e23)) & (e23 = op2(e23, e21))) | ((e22 = op2(e21, e23)) & (e22 = op2(e23, e21))) | ((e21 = op2(e21, e23)) & (e21 = op2(e23, e21))) | ((e20 = op2(e21, e23)) & (e20 = op2(e23, e21)))) & (e21 = op2(e23, e23))) & (e21 = op2(e23, e23))) | (~ ((((e23 = op2(e20, e23)) & (e23 = op2(e23, e20))) | ((e22 = op2(e20, e23)) & (e22 = op2(e23, e20))) | ((e21 = op2(e20, e23)) & (e21 = op2(e23, e20))) | ((e20 = op2(e20, e23)) & (e20 = op2(e23, e20)))) & (e20 = op2(e23, e23))) & (e20 = op2(e23, e23))) | (~ ((((e23 = op2(e23, e22)) & (e23 = op2(e22, e23))) | ((e22 = op2(e23, e22)) & (e22 = op2(e22, e23))) | ((e21 = op2(e23, e22)) & (e21 = op2(e22, e23))) | ((e20 = op2(e23, e22)) & (e20 = op2(e22, e23)))) & (e23 = op2(e22, e22))) & (e23 = op2(e22, e22))) | (~ ((((e23 = op2(e22, e22)) & (e23 = op2(e22, e22))) | ((e22 = op2(e22, e22)) & (e22 = op2(e22, e22))) | ((e21 = op2(e22, e22)) & (e21 = op2(e22, e22))) | ((e20 = op2(e22, e22)) & (e20 = op2(e22, e22)))) & (e22 = op2(e22, e22))) & (e22 = op2(e22, e22))) | (~ ((((e23 = op2(e21, e22)) & (e23 = op2(e22, e21))) | ((e22 = op2(e21, e22)) & (e22 = op2(e22, e21))) | ((e21 = op2(e21, e22)) & (e21 = op2(e22, e21))) | ((e20 = op2(e21, e22)) & (e20 = op2(e22, e21)))) & (e21 = op2(e22, e22))) & (e21 = op2(e22, e22))) | (~ ((((e23 = op2(e20, e22)) & (e23 = op2(e22, e20))) | ((e22 = op2(e20, e22)) & (e22 = op2(e22, e20))) | ((e21 = op2(e20, e22)) & (e21 = op2(e22, e20))) | ((e20 = op2(e20, e22)) & (e20 = op2(e22, e20)))) & (e20 = op2(e22, e22))) & (e20 = op2(e22, e22))) | (~ ((((e23 = op2(e23, e21)) & (e23 = op2(e21, e23))) | ((e22 = op2(e23, e21)) & (e22 = op2(e21, e23))) | ((e21 = op2(e23, e21)) & (e21 = op2(e21, e23))) | ((e20 = op2(e23, e21)) & (e20 = op2(e21, e23)))) & (e23 = op2(e21, e21))) & (e23 = op2(e21, e21))) | (~ ((((e23 = op2(e22, e21)) & (e23 = op2(e21, e22))) | ((e22 = op2(e22, e21)) & (e22 = op2(e21, e22))) | ((e21 = op2(e22, e21)) & (e21 = op2(e21, e22))) | ((e20 = op2(e22, e21)) & (e20 = op2(e21, e22)))) & (e22 = op2(e21, e21))) & (e22 = op2(e21, e21))) | (~ ((((e23 = op2(e21, e21)) & (e23 = op2(e21, e21))) | ((e22 = op2(e21, e21)) & (e22 = op2(e21, e21))) | ((e21 = op2(e21, e21)) & (e21 = op2(e21, e21))) | ((e20 = op2(e21, e21)) & (e20 = op2(e21, e21)))) & (e21 = op2(e21, e21))) & (e21 = op2(e21, e21))) | (~ ((((e23 = op2(e20, e21)) & (e23 = op2(e21, e20))) | ((e22 = op2(e20, e21)) & (e22 = op2(e21, e20))) | ((e21 = op2(e20, e21)) & (e21 = op2(e21, e20))) | ((e20 = op2(e20, e21)) & (e20 = op2(e21, e20)))) & (e20 = op2(e21, e21))) & (e20 = op2(e21, e21))) | (~ ((((e23 = op2(e23, e20)) & (e23 = op2(e20, e23))) | ((e22 = op2(e23, e20)) & (e22 = op2(e20, e23))) | ((e21 = op2(e23, e20)) & (e21 = op2(e20, e23))) | ((e20 = op2(e23, e20)) & (e20 = op2(e20, e23)))) & (op2(e20, e20) = e23)) & (op2(e20, e20) = e23)) | (~ ((((e23 = op2(e22, e20)) & (e23 = op2(e20, e22))) | ((e22 = op2(e22, e20)) & (e22 = op2(e20, e22))) | ((e21 = op2(e22, e20)) & (e21 = op2(e20, e22))) | ((e20 = op2(e22, e20)) & (e20 = op2(e20, e22)))) & (op2(e20, e20) = e22)) & (op2(e20, e20) = e22)) | (~ ((((e23 = op2(e21, e20)) & (e23 = op2(e20, e21))) | ((e22 = op2(e21, e20)) & (e22 = op2(e20, e21))) | ((e21 = op2(e21, e20)) & (e21 = op2(e20, e21))) | ((e20 = op2(e21, e20)) & (e20 = op2(e20, e21)))) & (op2(e20, e20) = e21)) & (op2(e20, e20) = e21)) | (~ ((((op2(e20, e20) = e23) & (op2(e20, e20) = e23)) | ((op2(e20, e20) = e22) & (op2(e20, e20) = e22)) | ((op2(e20, e20) = e21) & (op2(e20, e20) = e21)) | ((e20 = op2(e20, e20)) & (e20 = op2(e20, e20)))) & (e20 = op2(e20, e20))) & (e20 = op2(e20, e20)))) & (((e23 = op2(e23, op2(e23, e23))) & ~ (op2(e23, e23) = op2(e23, e23)) & (op2(e23, e23) = op2(e23, e23))) | ((e22 = op2(e22, op2(e23, e22))) & ~ (op2(e22, e22) = op2(e23, e22)) & (op2(e22, e23) = op2(e23, e22))) | ((e21 = op2(e21, op2(e23, e21))) & ~ (op2(e21, e21) = op2(e23, e21)) & (op2(e21, e23) = op2(e23, e21))) | ((e20 = op2(e20, op2(e23, e20))) & ~ (op2(e20, e20) = op2(e23, e20)) & (op2(e20, e23) = op2(e23, e20))) | ((e23 = op2(e23, op2(e22, e23))) & ~ (op2(e22, e23) = op2(e23, e23)) & (op2(e22, e23) = op2(e23, e22))) | ((e22 = op2(e22, op2(e22, e22))) & ~ (op2(e22, e22) = op2(e22, e22)) & (op2(e22, e22) = op2(e22, e22))) | ((e21 = op2(e21, op2(e22, e21))) & ~ (op2(e21, e21) = op2(e22, e21)) & (op2(e21, e22) = op2(e22, e21))) | ((e20 = op2(e20, op2(e22, e20))) & ~ (op2(e20, e20) = op2(e22, e20)) & (op2(e20, e22) = op2(e22, e20))) | ((e23 = op2(e23, op2(e21, e23))) & ~ (op2(e21, e23) = op2(e23, e23)) & (op2(e21, e23) = op2(e23, e21))) | ((e22 = op2(e22, op2(e21, e22))) & ~ (op2(e21, e22) = op2(e22, e22)) & (op2(e21, e22) = op2(e22, e21))) | ((e21 = op2(e21, op2(e21, e21))) & ~ (op2(e21, e21) = op2(e21, e21)) & (op2(e21, e21) = op2(e21, e21))) | ((e20 = op2(e20, op2(e21, e20))) & ~ (op2(e20, e20) = op2(e21, e20)) & (op2(e20, e21) = op2(e21, e20))) | ((e23 = op2(e23, op2(e20, e23))) & ~ (op2(e20, e23) = op2(e23, e23)) & (op2(e20, e23) = op2(e23, e20))) | ((e22 = op2(e22, op2(e20, e22))) & ~ (op2(e20, e22) = op2(e22, e22)) & (op2(e20, e22) = op2(e22, e20))) | ((e21 = op2(e21, op2(e20, e21))) & ~ (op2(e20, e21) = op2(e21, e21)) & (op2(e20, e21) = op2(e21, e20))) | ((e20 = op2(e20, op2(e20, e20))) & ~ (op2(e20, e20) = op2(e20, e20)) & (op2(e20, e20) = op2(e20, e20)))) & ((~ (e23 = op2(e23, e23)) & ~ (e22 = op2(e22, e23)) & ~ (e21 = op2(e21, e23)) & ~ (e20 = op2(e20, e23)) & (e23 = op2(e23, e23)) & (e23 = op2(e23, e23))) | (~ (e23 = op2(e23, e23)) & ~ (e22 = op2(e22, e23)) & ~ (e21 = op2(e21, e23)) & ~ (e20 = op2(e20, e23)) & (e23 = op2(e22, e23)) & (e23 = op2(e23, e22))) | (~ (e23 = op2(e23, e23)) & ~ (e22 = op2(e22, e23)) & ~ (e21 = op2(e21, e23)) & ~ (e20 = op2(e20, e23)) & (e23 = op2(e21, e23)) & (e23 = op2(e23, e21))) | (~ (e23 = op2(e23, e23)) & ~ (e22 = op2(e22, e23)) & ~ (e21 = op2(e21, e23)) & ~ (e20 = op2(e20, e23)) & (e23 = op2(e20, e23)) & (e23 = op2(e23, e20))) | (~ (e23 = op2(e23, e22)) & ~ (e22 = op2(e22, e22)) & ~ (e21 = op2(e21, e22)) & ~ (e20 = op2(e20, e22)) & (e22 = op2(e23, e22)) & (e22 = op2(e22, e23))) | (~ (e23 = op2(e23, e22)) & ~ (e22 = op2(e22, e22)) & ~ (e21 = op2(e21, e22)) & ~ (e20 = op2(e20, e22)) & (e22 = op2(e22, e22)) & (e22 = op2(e22, e22))) | (~ (e23 = op2(e23, e22)) & ~ (e22 = op2(e22, e22)) & ~ (e21 = op2(e21, e22)) & ~ (e20 = op2(e20, e22)) & (e22 = op2(e21, e22)) & (e22 = op2(e22, e21))) | (~ (e23 = op2(e23, e22)) & ~ (e22 = op2(e22, e22)) & ~ (e21 = op2(e21, e22)) & ~ (e20 = op2(e20, e22)) & (e22 = op2(e20, e22)) & (e22 = op2(e22, e20))) | (~ (e23 = op2(e23, e21)) & ~ (e22 = op2(e22, e21)) & ~ (e21 = op2(e21, e21)) & ~ (e20 = op2(e20, e21)) & (e21 = op2(e23, e21)) & (e21 = op2(e21, e23))) | (~ (e23 = op2(e23, e21)) & ~ (e22 = op2(e22, e21)) & ~ (e21 = op2(e21, e21)) & ~ (e20 = op2(e20, e21)) & (e21 = op2(e22, e21)) & (e21 = op2(e21, e22))) | (~ (e23 = op2(e23, e21)) & ~ (e22 = op2(e22, e21)) & ~ (e21 = op2(e21, e21)) & ~ (e20 = op2(e20, e21)) & (e21 = op2(e21, e21)) & (e21 = op2(e21, e21))) | (~ (e23 = op2(e23, e21)) & ~ (e22 = op2(e22, e21)) & ~ (e21 = op2(e21, e21)) & ~ (e20 = op2(e20, e21)) & (e21 = op2(e20, e21)) & (e21 = op2(e21, e20))) | (~ (e23 = op2(e23, e20)) & ~ (e22 = op2(e22, e20)) & ~ (e21 = op2(e21, e20)) & ~ (e20 = op2(e20, e20)) & (e20 = op2(e23, e20)) & (e20 = op2(e20, e23))) | (~ (e23 = op2(e23, e20)) & ~ (e22 = op2(e22, e20)) & ~ (e21 = op2(e21, e20)) & ~ (e20 = op2(e20, e20)) & (e20 = op2(e22, e20)) & (e20 = op2(e20, e22))) | (~ (e23 = op2(e23, e20)) & ~ (e22 = op2(e22, e20)) & ~ (e21 = op2(e21, e20)) & ~ (e20 = op2(e20, e20)) & (e20 = op2(e21, e20)) & (e20 = op2(e20, e21))) | (~ (e23 = op2(e23, e20)) & ~ (e22 = op2(e22, e20)) & ~ (e21 = op2(e21, e20)) & ~ (e20 = op2(e20, e20)) & (e20 = op2(e20, e20)) & (e20 = op2(e20, e20))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG109+1.p', ax11)).
fof(f2581, plain, (spl102_261 | spl102_260 | spl102_259 | spl102_258 | spl102_257 | spl102_256 | spl102_255 | spl102_254 | spl102_253 | spl102_252 | spl102_251 | spl102_250 | spl102_249 | spl102_248 | spl102_247 | ~ spl102_113), inference(avatar_split_clause, [], [f886, f1533, f2418, f2428, f2438, f2448, f2458, f2468, f2478, f2488, f2498, f2508, f2518, f2528, f2538, f2548, f2558])).
fof(f886, plain, (~ (e20 = op2(e20, e23)) | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45), inference(cnf_transformation, [], [f114])).
fof(f2578, plain, (spl102_261 | spl102_260 | spl102_259 | spl102_258 | spl102_257 | spl102_256 | spl102_255 | spl102_254 | spl102_253 | spl102_252 | spl102_251 | spl102_250 | spl102_249 | spl102_248 | spl102_247 | ~ spl102_68), inference(avatar_split_clause, [], [f889, f1341, f2418, f2428, f2438, f2448, f2458, f2468, f2478, f2488, f2498, f2508, f2518, f2528, f2538, f2548, f2558])).
fof(f889, plain, (~ (e23 = op2(e23, e23)) | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45), inference(cnf_transformation, [], [f114])).
fof(f2577, plain, (spl102_245 | spl102_243 | spl102_241 | spl102_239 | spl102_236 | spl102_234 | spl102_232 | spl102_230 | spl102_227 | spl102_224 | spl102_222 | spl102_220 | spl102_217 | spl102_214 | spl102_211), inference(avatar_split_clause, [], [f984, f2244, f2258, f2272, f2286, f2296, f2306, f2320, f2334, f2344, f2354, f2364, f2378, f2388, f2398, f2408])).
fof(f2408, plain, (spl102_245 <=> sP60), introduced(avatar_definition, [new_symbols(naming, [spl102_245])])).
fof(f2398, plain, (spl102_243 <=> sP61), introduced(avatar_definition, [new_symbols(naming, [spl102_243])])).
fof(f2388, plain, (spl102_241 <=> sP62), introduced(avatar_definition, [new_symbols(naming, [spl102_241])])).
fof(f2378, plain, (spl102_239 <=> sP63), introduced(avatar_definition, [new_symbols(naming, [spl102_239])])).
fof(f2364, plain, (spl102_236 <=> sP64), introduced(avatar_definition, [new_symbols(naming, [spl102_236])])).
fof(f2354, plain, (spl102_234 <=> sP65), introduced(avatar_definition, [new_symbols(naming, [spl102_234])])).
fof(f2344, plain, (spl102_232 <=> sP66), introduced(avatar_definition, [new_symbols(naming, [spl102_232])])).
fof(f2334, plain, (spl102_230 <=> sP67), introduced(avatar_definition, [new_symbols(naming, [spl102_230])])).
fof(f2320, plain, (spl102_227 <=> sP68), introduced(avatar_definition, [new_symbols(naming, [spl102_227])])).
fof(f2306, plain, (spl102_224 <=> sP69), introduced(avatar_definition, [new_symbols(naming, [spl102_224])])).
fof(f2296, plain, (spl102_222 <=> sP70), introduced(avatar_definition, [new_symbols(naming, [spl102_222])])).
fof(f2286, plain, (spl102_220 <=> sP71), introduced(avatar_definition, [new_symbols(naming, [spl102_220])])).
fof(f2272, plain, (spl102_217 <=> sP72), introduced(avatar_definition, [new_symbols(naming, [spl102_217])])).
fof(f2258, plain, (spl102_214 <=> sP73), introduced(avatar_definition, [new_symbols(naming, [spl102_214])])).
fof(f2244, plain, (spl102_211 <=> sP74), introduced(avatar_definition, [new_symbols(naming, [spl102_211])])).
fof(f984, plain, (sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60), inference(trivial_inequality_removal, [], [f891])).
fof(f891, plain, (~ (op2(e23, e23) = op2(e23, e23)) | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60), inference(cnf_transformation, [], [f114])).
fof(f2571, plain, (spl102_210 | spl102_209 | spl102_208 | spl102_207 | spl102_206 | spl102_205 | spl102_204 | spl102_203 | spl102_202 | spl102_201 | spl102_200 | spl102_199 | spl102_198 | spl102_197 | spl102_196 | spl102_68), inference(avatar_split_clause, [], [f893, f1341, f2109, f2118, f2127, f2136, f2145, f2154, f2163, f2172, f2181, f2190, f2199, f2208, f2217, f2226, f2235])).
fof(f2235, plain, (spl102_210 <=> sP75), introduced(avatar_definition, [new_symbols(naming, [spl102_210])])).
fof(f2226, plain, (spl102_209 <=> sP76), introduced(avatar_definition, [new_symbols(naming, [spl102_209])])).
fof(f2217, plain, (spl102_208 <=> sP77), introduced(avatar_definition, [new_symbols(naming, [spl102_208])])).
fof(f2208, plain, (spl102_207 <=> sP78), introduced(avatar_definition, [new_symbols(naming, [spl102_207])])).
fof(f2199, plain, (spl102_206 <=> sP79), introduced(avatar_definition, [new_symbols(naming, [spl102_206])])).
fof(f2190, plain, (spl102_205 <=> sP80), introduced(avatar_definition, [new_symbols(naming, [spl102_205])])).
fof(f2181, plain, (spl102_204 <=> sP81), introduced(avatar_definition, [new_symbols(naming, [spl102_204])])).
fof(f2172, plain, (spl102_203 <=> sP82), introduced(avatar_definition, [new_symbols(naming, [spl102_203])])).
fof(f2163, plain, (spl102_202 <=> sP83), introduced(avatar_definition, [new_symbols(naming, [spl102_202])])).
fof(f2154, plain, (spl102_201 <=> sP84), introduced(avatar_definition, [new_symbols(naming, [spl102_201])])).
fof(f2145, plain, (spl102_200 <=> sP85), introduced(avatar_definition, [new_symbols(naming, [spl102_200])])).
fof(f2136, plain, (spl102_199 <=> sP86), introduced(avatar_definition, [new_symbols(naming, [spl102_199])])).
fof(f2127, plain, (spl102_198 <=> sP87), introduced(avatar_definition, [new_symbols(naming, [spl102_198])])).
fof(f2118, plain, (spl102_197 <=> sP88), introduced(avatar_definition, [new_symbols(naming, [spl102_197])])).
fof(f2109, plain, (spl102_196 <=> sP89), introduced(avatar_definition, [new_symbols(naming, [spl102_196])])).
fof(f893, plain, ((e23 = op2(e23, e23)) | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75), inference(cnf_transformation, [], [f114])).
fof(f2567, plain, (spl102_210 | spl102_209 | spl102_208 | spl102_207 | spl102_206 | spl102_205 | spl102_204 | spl102_203 | spl102_202 | spl102_201 | spl102_200 | spl102_199 | spl102_198 | spl102_197 | spl102_196 | ~ spl102_68), inference(avatar_split_clause, [], [f988, f1341, f2109, f2118, f2127, f2136, f2145, f2154, f2163, f2172, f2181, f2190, f2199, f2208, f2217, f2226, f2235])).
fof(f988, plain, (~ (e23 = op2(e23, e23)) | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75), inference(duplicate_literal_removal, [], [f897])).
fof(f897, plain, (~ (e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23)) | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75), inference(cnf_transformation, [], [f114])).
fof(f2566, plain, (~ spl102_261 | spl102_125), inference(avatar_split_clause, [], [f878, f1584, f2558])).
fof(f878, plain, ((e20 = op2(e20, e20)) | ~ sP45), inference(cnf_transformation, [], [f217])).
fof(f217, plain, ((~ (e23 = op2(e23, e20)) & ~ (e22 = op2(e22, e20)) & ~ (e21 = op2(e21, e20)) & ~ (e20 = op2(e20, e20)) & (e20 = op2(e20, e20)) & (e20 = op2(e20, e20))) | ~ sP45), inference(nnf_transformation, [], [f69])).
fof(f2556, plain, (~ spl102_260 | spl102_121), inference(avatar_split_clause, [], [f872, f1567, f2548])).
fof(f872, plain, ((e20 = op2(e20, e21)) | ~ sP46), inference(cnf_transformation, [], [f216])).
fof(f216, plain, ((~ (e23 = op2(e23, e20)) & ~ (e22 = op2(e22, e20)) & ~ (e21 = op2(e21, e20)) & ~ (e20 = op2(e20, e20)) & (e20 = op2(e21, e20)) & (e20 = op2(e20, e21))) | ~ sP46), inference(nnf_transformation, [], [f70])).
fof(f2546, plain, (~ spl102_259 | spl102_117), inference(avatar_split_clause, [], [f866, f1550, f2538])).
fof(f866, plain, ((e20 = op2(e20, e22)) | ~ sP47), inference(cnf_transformation, [], [f215])).
fof(f215, plain, ((~ (e23 = op2(e23, e20)) & ~ (e22 = op2(e22, e20)) & ~ (e21 = op2(e21, e20)) & ~ (e20 = op2(e20, e20)) & (e20 = op2(e22, e20)) & (e20 = op2(e20, e22))) | ~ sP47), inference(nnf_transformation, [], [f71])).
fof(f2545, plain, (~ spl102_259 | spl102_93), inference(avatar_split_clause, [], [f867, f1448, f2538])).
fof(f867, plain, ((e20 = op2(e22, e20)) | ~ sP47), inference(cnf_transformation, [], [f215])).
fof(f2535, plain, (~ spl102_258 | spl102_77), inference(avatar_split_clause, [], [f861, f1380, f2528])).
fof(f861, plain, ((e20 = op2(e23, e20)) | ~ sP48), inference(cnf_transformation, [], [f214])).
fof(f214, plain, ((~ (e23 = op2(e23, e20)) & ~ (e22 = op2(e22, e20)) & ~ (e21 = op2(e21, e20)) & ~ (e20 = op2(e20, e20)) & (e20 = op2(e23, e20)) & (e20 = op2(e20, e23))) | ~ sP48), inference(nnf_transformation, [], [f72])).
fof(f2526, plain, (~ spl102_257 | spl102_110), inference(avatar_split_clause, [], [f854, f1520, f2518])).
fof(f854, plain, ((e21 = op2(e21, e20)) | ~ sP49), inference(cnf_transformation, [], [f213])).
fof(f213, plain, ((~ (e23 = op2(e23, e21)) & ~ (e22 = op2(e22, e21)) & ~ (e21 = op2(e21, e21)) & ~ (e20 = op2(e20, e21)) & (e21 = op2(e20, e21)) & (e21 = op2(e21, e20))) | ~ sP49), inference(nnf_transformation, [], [f73])).
fof(f2516, plain, (~ spl102_256 | spl102_106), inference(avatar_split_clause, [], [f848, f1503, f2508])).
fof(f848, plain, ((e21 = op2(e21, e21)) | ~ sP50), inference(cnf_transformation, [], [f212])).
fof(f212, plain, ((~ (e23 = op2(e23, e21)) & ~ (e22 = op2(e22, e21)) & ~ (e21 = op2(e21, e21)) & ~ (e20 = op2(e20, e21)) & (e21 = op2(e21, e21)) & (e21 = op2(e21, e21))) | ~ sP50), inference(nnf_transformation, [], [f74])).
fof(f2513, plain, (~ spl102_256 | ~ spl102_106), inference(avatar_split_clause, [], [f851, f1503, f2508])).
fof(f851, plain, (~ (e21 = op2(e21, e21)) | ~ sP50), inference(cnf_transformation, [], [f212])).
fof(f2512, plain, (~ spl102_256 | ~ spl102_91), inference(avatar_split_clause, [], [f852, f1439, f2508])).
fof(f852, plain, (~ (e22 = op2(e22, e21)) | ~ sP50), inference(cnf_transformation, [], [f212])).
fof(f2506, plain, (~ spl102_255 | spl102_102), inference(avatar_split_clause, [], [f842, f1486, f2498])).
fof(f842, plain, ((e21 = op2(e21, e22)) | ~ sP51), inference(cnf_transformation, [], [f211])).
fof(f211, plain, ((~ (e23 = op2(e23, e21)) & ~ (e22 = op2(e22, e21)) & ~ (e21 = op2(e21, e21)) & ~ (e20 = op2(e20, e21)) & (e21 = op2(e22, e21)) & (e21 = op2(e21, e22))) | ~ sP51), inference(nnf_transformation, [], [f75])).
fof(f2505, plain, (~ spl102_255 | spl102_90), inference(avatar_split_clause, [], [f843, f1435, f2498])).
fof(f843, plain, ((e21 = op2(e22, e21)) | ~ sP51), inference(cnf_transformation, [], [f211])).
fof(f2502, plain, (~ spl102_255 | ~ spl102_91), inference(avatar_split_clause, [], [f846, f1439, f2498])).
fof(f846, plain, (~ (e22 = op2(e22, e21)) | ~ sP51), inference(cnf_transformation, [], [f211])).
fof(f2496, plain, (~ spl102_254 | spl102_98), inference(avatar_split_clause, [], [f836, f1469, f2488])).
fof(f836, plain, ((e21 = op2(e21, e23)) | ~ sP52), inference(cnf_transformation, [], [f210])).
fof(f210, plain, ((~ (e23 = op2(e23, e21)) & ~ (e22 = op2(e22, e21)) & ~ (e21 = op2(e21, e21)) & ~ (e20 = op2(e20, e21)) & (e21 = op2(e23, e21)) & (e21 = op2(e21, e23))) | ~ sP52), inference(nnf_transformation, [], [f76])).
fof(f2495, plain, (~ spl102_254 | spl102_74), inference(avatar_split_clause, [], [f837, f1367, f2488])).
fof(f837, plain, ((e21 = op2(e23, e21)) | ~ sP52), inference(cnf_transformation, [], [f210])).
fof(f2493, plain, (~ spl102_254 | ~ spl102_106), inference(avatar_split_clause, [], [f839, f1503, f2488])).
fof(f839, plain, (~ (e21 = op2(e21, e21)) | ~ sP52), inference(cnf_transformation, [], [f210])).
fof(f2492, plain, (~ spl102_254 | ~ spl102_91), inference(avatar_split_clause, [], [f840, f1439, f2488])).
fof(f840, plain, (~ (e22 = op2(e22, e21)) | ~ sP52), inference(cnf_transformation, [], [f210])).
fof(f2486, plain, (~ spl102_253 | spl102_95), inference(avatar_split_clause, [], [f830, f1456, f2478])).
fof(f830, plain, ((e22 = op2(e22, e20)) | ~ sP53), inference(cnf_transformation, [], [f209])).
fof(f209, plain, ((~ (e23 = op2(e23, e22)) & ~ (e22 = op2(e22, e22)) & ~ (e21 = op2(e21, e22)) & ~ (e20 = op2(e20, e22)) & (e22 = op2(e20, e22)) & (e22 = op2(e22, e20))) | ~ sP53), inference(nnf_transformation, [], [f77])).
fof(f2475, plain, (~ spl102_252 | spl102_103), inference(avatar_split_clause, [], [f825, f1490, f2468])).
fof(f825, plain, ((e22 = op2(e21, e22)) | ~ sP54), inference(cnf_transformation, [], [f208])).
fof(f208, plain, ((~ (e23 = op2(e23, e22)) & ~ (e22 = op2(e22, e22)) & ~ (e21 = op2(e21, e22)) & ~ (e20 = op2(e20, e22)) & (e22 = op2(e21, e22)) & (e22 = op2(e22, e21))) | ~ sP54), inference(nnf_transformation, [], [f78])).
fof(f2474, plain, (~ spl102_252 | ~ spl102_117), inference(avatar_split_clause, [], [f826, f1550, f2468])).
fof(f826, plain, (~ (e20 = op2(e20, e22)) | ~ sP54), inference(cnf_transformation, [], [f208])).
fof(f2473, plain, (~ spl102_252 | ~ spl102_102), inference(avatar_split_clause, [], [f827, f1486, f2468])).
fof(f827, plain, (~ (e21 = op2(e21, e22)) | ~ sP54), inference(cnf_transformation, [], [f208])).
fof(f2471, plain, (~ spl102_252 | ~ spl102_72), inference(avatar_split_clause, [], [f829, f1358, f2468])).
fof(f829, plain, (~ (e23 = op2(e23, e22)) | ~ sP54), inference(cnf_transformation, [], [f208])).
fof(f2466, plain, (~ spl102_251 | spl102_87), inference(avatar_split_clause, [], [f818, f1422, f2458])).
fof(f818, plain, ((e22 = op2(e22, e22)) | ~ sP55), inference(cnf_transformation, [], [f207])).
fof(f207, plain, ((~ (e23 = op2(e23, e22)) & ~ (e22 = op2(e22, e22)) & ~ (e21 = op2(e21, e22)) & ~ (e20 = op2(e20, e22)) & (e22 = op2(e22, e22)) & (e22 = op2(e22, e22))) | ~ sP55), inference(nnf_transformation, [], [f79])).
fof(f2464, plain, (~ spl102_251 | ~ spl102_117), inference(avatar_split_clause, [], [f820, f1550, f2458])).
fof(f820, plain, (~ (e20 = op2(e20, e22)) | ~ sP55), inference(cnf_transformation, [], [f207])).
fof(f2462, plain, (~ spl102_251 | ~ spl102_87), inference(avatar_split_clause, [], [f822, f1422, f2458])).
fof(f822, plain, (~ (e22 = op2(e22, e22)) | ~ sP55), inference(cnf_transformation, [], [f207])).
fof(f2455, plain, (~ spl102_250 | spl102_71), inference(avatar_split_clause, [], [f813, f1354, f2448])).
fof(f813, plain, ((e22 = op2(e23, e22)) | ~ sP56), inference(cnf_transformation, [], [f206])).
fof(f206, plain, ((~ (e23 = op2(e23, e22)) & ~ (e22 = op2(e22, e22)) & ~ (e21 = op2(e21, e22)) & ~ (e20 = op2(e20, e22)) & (e22 = op2(e23, e22)) & (e22 = op2(e22, e23))) | ~ sP56), inference(nnf_transformation, [], [f80])).
fof(f2446, plain, (~ spl102_249 | spl102_80), inference(avatar_split_clause, [], [f806, f1392, f2438])).
fof(f806, plain, ((e23 = op2(e23, e20)) | ~ sP57), inference(cnf_transformation, [], [f205])).
fof(f205, plain, ((~ (e23 = op2(e23, e23)) & ~ (e22 = op2(e22, e23)) & ~ (e21 = op2(e21, e23)) & ~ (e20 = op2(e20, e23)) & (e23 = op2(e20, e23)) & (e23 = op2(e23, e20))) | ~ sP57), inference(nnf_transformation, [], [f81])).
fof(f2436, plain, (~ spl102_248 | spl102_76), inference(avatar_split_clause, [], [f800, f1375, f2428])).
fof(f800, plain, ((e23 = op2(e23, e21)) | ~ sP58), inference(cnf_transformation, [], [f204])).
fof(f204, plain, ((~ (e23 = op2(e23, e23)) & ~ (e22 = op2(e22, e23)) & ~ (e21 = op2(e21, e23)) & ~ (e20 = op2(e20, e23)) & (e23 = op2(e21, e23)) & (e23 = op2(e23, e21))) | ~ sP58), inference(nnf_transformation, [], [f82])).
fof(f2425, plain, (~ spl102_247 | spl102_84), inference(avatar_split_clause, [], [f795, f1409, f2418])).
fof(f795, plain, ((e23 = op2(e22, e23)) | ~ sP59), inference(cnf_transformation, [], [f203])).
fof(f203, plain, ((~ (e23 = op2(e23, e23)) & ~ (e22 = op2(e22, e23)) & ~ (e21 = op2(e21, e23)) & ~ (e20 = op2(e20, e23)) & (e23 = op2(e22, e23)) & (e23 = op2(e23, e22))) | ~ sP59), inference(nnf_transformation, [], [f83])).
fof(f2424, plain, (~ spl102_247 | ~ spl102_113), inference(avatar_split_clause, [], [f796, f1533, f2418])).
fof(f796, plain, (~ (e20 = op2(e20, e23)) | ~ sP59), inference(cnf_transformation, [], [f203])).
fof(f2421, plain, (~ spl102_247 | ~ spl102_68), inference(avatar_split_clause, [], [f799, f1341, f2418])).
fof(f799, plain, (~ (e23 = op2(e23, e23)) | ~ sP59), inference(cnf_transformation, [], [f203])).
fof(f2416, plain, ~ spl102_245, inference(avatar_split_clause, [], [f989, f2408])).
fof(f989, plain, ~ sP60, inference(trivial_inequality_removal, [], [f792])).
fof(f792, plain, (~ (op2(e20, e20) = op2(e20, e20)) | ~ sP60), inference(cnf_transformation, [], [f202])).
fof(f202, plain, (((e20 = op2(e20, op2(e20, e20))) & ~ (op2(e20, e20) = op2(e20, e20)) & (op2(e20, e20) = op2(e20, e20))) | ~ sP60), inference(nnf_transformation, [], [f84])).
fof(f2406, plain, (~ spl102_243 | spl102_238), inference(avatar_split_clause, [], [f788, f2373, f2398])).
fof(f788, plain, ((op2(e20, e21) = op2(e21, e20)) | ~ sP61), inference(cnf_transformation, [], [f201])).
fof(f201, plain, (((e21 = op2(e21, op2(e20, e21))) & ~ (op2(e20, e21) = op2(e21, e21)) & (op2(e20, e21) = op2(e21, e20))) | ~ sP61), inference(nnf_transformation, [], [f85])).
fof(f2396, plain, (~ spl102_241 | spl102_229), inference(avatar_split_clause, [], [f785, f2329, f2388])).
fof(f785, plain, ((op2(e20, e22) = op2(e22, e20)) | ~ sP62), inference(cnf_transformation, [], [f200])).
fof(f200, plain, (((e22 = op2(e22, op2(e20, e22))) & ~ (op2(e20, e22) = op2(e22, e22)) & (op2(e20, e22) = op2(e22, e20))) | ~ sP62), inference(nnf_transformation, [], [f86])).
fof(f2386, plain, (~ spl102_239 | spl102_219), inference(avatar_split_clause, [], [f782, f2281, f2378])).
fof(f782, plain, ((op2(e20, e23) = op2(e23, e20)) | ~ sP63), inference(cnf_transformation, [], [f199])).
fof(f199, plain, (((e23 = op2(e23, op2(e20, e23))) & ~ (op2(e20, e23) = op2(e23, e23)) & (op2(e20, e23) = op2(e23, e20))) | ~ sP63), inference(nnf_transformation, [], [f87])).
fof(f2376, plain, (~ spl102_236 | spl102_238), inference(avatar_split_clause, [], [f779, f2373, f2364])).
fof(f779, plain, ((op2(e20, e21) = op2(e21, e20)) | ~ sP64), inference(cnf_transformation, [], [f198])).
fof(f198, plain, (((e20 = op2(e20, op2(e21, e20))) & ~ (op2(e20, e20) = op2(e21, e20)) & (op2(e20, e21) = op2(e21, e20))) | ~ sP64), inference(nnf_transformation, [], [f88])).
fof(f2371, plain, (~ spl102_236 | spl102_237), inference(avatar_split_clause, [], [f781, f2368, f2364])).
fof(f781, plain, ((e20 = op2(e20, op2(e21, e20))) | ~ sP64), inference(cnf_transformation, [], [f198])).
fof(f2362, plain, ~ spl102_234, inference(avatar_split_clause, [], [f990, f2354])).
fof(f990, plain, ~ sP65, inference(trivial_inequality_removal, [], [f777])).
fof(f777, plain, (~ (op2(e21, e21) = op2(e21, e21)) | ~ sP65), inference(cnf_transformation, [], [f197])).
fof(f197, plain, (((e21 = op2(e21, op2(e21, e21))) & ~ (op2(e21, e21) = op2(e21, e21)) & (op2(e21, e21) = op2(e21, e21))) | ~ sP65), inference(nnf_transformation, [], [f89])).
fof(f2352, plain, (~ spl102_232 | spl102_226), inference(avatar_split_clause, [], [f773, f2315, f2344])).
fof(f773, plain, ((op2(e21, e22) = op2(e22, e21)) | ~ sP66), inference(cnf_transformation, [], [f196])).
fof(f196, plain, (((e22 = op2(e22, op2(e21, e22))) & ~ (op2(e21, e22) = op2(e22, e22)) & (op2(e21, e22) = op2(e22, e21))) | ~ sP66), inference(nnf_transformation, [], [f90])).
fof(f2351, plain, (~ spl102_232 | spl102_233), inference(avatar_split_clause, [], [f775, f2348, f2344])).
fof(f775, plain, ((e22 = op2(e22, op2(e21, e22))) | ~ sP66), inference(cnf_transformation, [], [f196])).
fof(f2342, plain, (~ spl102_230 | spl102_216), inference(avatar_split_clause, [], [f770, f2267, f2334])).
fof(f770, plain, ((op2(e21, e23) = op2(e23, e21)) | ~ sP67), inference(cnf_transformation, [], [f195])).
fof(f195, plain, (((e23 = op2(e23, op2(e21, e23))) & ~ (op2(e21, e23) = op2(e23, e23)) & (op2(e21, e23) = op2(e23, e21))) | ~ sP67), inference(nnf_transformation, [], [f91])).
fof(f2341, plain, (~ spl102_230 | spl102_231), inference(avatar_split_clause, [], [f772, f2338, f2334])).
fof(f772, plain, ((e23 = op2(e23, op2(e21, e23))) | ~ sP67), inference(cnf_transformation, [], [f195])).
fof(f2332, plain, (~ spl102_227 | spl102_229), inference(avatar_split_clause, [], [f767, f2329, f2320])).
fof(f767, plain, ((op2(e20, e22) = op2(e22, e20)) | ~ sP68), inference(cnf_transformation, [], [f194])).
fof(f194, plain, (((e20 = op2(e20, op2(e22, e20))) & ~ (op2(e20, e20) = op2(e22, e20)) & (op2(e20, e22) = op2(e22, e20))) | ~ sP68), inference(nnf_transformation, [], [f92])).
fof(f2318, plain, (~ spl102_224 | spl102_226), inference(avatar_split_clause, [], [f764, f2315, f2306])).
fof(f764, plain, ((op2(e21, e22) = op2(e22, e21)) | ~ sP69), inference(cnf_transformation, [], [f193])).
fof(f193, plain, (((e21 = op2(e21, op2(e22, e21))) & ~ (op2(e21, e21) = op2(e22, e21)) & (op2(e21, e22) = op2(e22, e21))) | ~ sP69), inference(nnf_transformation, [], [f93])).
fof(f2313, plain, (~ spl102_224 | spl102_225), inference(avatar_split_clause, [], [f766, f2310, f2306])).
fof(f766, plain, ((e21 = op2(e21, op2(e22, e21))) | ~ sP69), inference(cnf_transformation, [], [f193])).
fof(f2304, plain, ~ spl102_222, inference(avatar_split_clause, [], [f991, f2296])).
fof(f991, plain, ~ sP70, inference(trivial_inequality_removal, [], [f762])).
fof(f762, plain, (~ (op2(e22, e22) = op2(e22, e22)) | ~ sP70), inference(cnf_transformation, [], [f192])).
fof(f192, plain, (((e22 = op2(e22, op2(e22, e22))) & ~ (op2(e22, e22) = op2(e22, e22)) & (op2(e22, e22) = op2(e22, e22))) | ~ sP70), inference(nnf_transformation, [], [f94])).
fof(f2294, plain, (~ spl102_220 | spl102_213), inference(avatar_split_clause, [], [f758, f2253, f2286])).
fof(f758, plain, ((op2(e22, e23) = op2(e23, e22)) | ~ sP71), inference(cnf_transformation, [], [f191])).
fof(f191, plain, (((e23 = op2(e23, op2(e22, e23))) & ~ (op2(e22, e23) = op2(e23, e23)) & (op2(e22, e23) = op2(e23, e22))) | ~ sP71), inference(nnf_transformation, [], [f95])).
fof(f2279, plain, (~ spl102_217 | spl102_218), inference(avatar_split_clause, [], [f757, f2276, f2272])).
fof(f757, plain, ((e20 = op2(e20, op2(e23, e20))) | ~ sP72), inference(cnf_transformation, [], [f190])).
fof(f190, plain, (((e20 = op2(e20, op2(e23, e20))) & ~ (op2(e20, e20) = op2(e23, e20)) & (op2(e20, e23) = op2(e23, e20))) | ~ sP72), inference(nnf_transformation, [], [f96])).
fof(f2265, plain, (~ spl102_214 | spl102_215), inference(avatar_split_clause, [], [f754, f2262, f2258])).
fof(f754, plain, ((e21 = op2(e21, op2(e23, e21))) | ~ sP73), inference(cnf_transformation, [], [f189])).
fof(f189, plain, (((e21 = op2(e21, op2(e23, e21))) & ~ (op2(e21, e21) = op2(e23, e21)) & (op2(e21, e23) = op2(e23, e21))) | ~ sP73), inference(nnf_transformation, [], [f97])).
fof(f2256, plain, (~ spl102_211 | spl102_213), inference(avatar_split_clause, [], [f749, f2253, f2244])).
fof(f749, plain, ((op2(e22, e23) = op2(e23, e22)) | ~ sP74), inference(cnf_transformation, [], [f188])).
fof(f188, plain, (((e22 = op2(e22, op2(e23, e22))) & ~ (op2(e22, e22) = op2(e23, e22)) & (op2(e22, e23) = op2(e23, e22))) | ~ sP74), inference(nnf_transformation, [], [f98])).
fof(f2242, plain, (~ spl102_210 | spl102_125), inference(avatar_split_clause, [], [f744, f1584, f2235])).
fof(f744, plain, ((e20 = op2(e20, e20)) | ~ sP75), inference(cnf_transformation, [], [f187])).
fof(f187, plain, (((((~ (op2(e20, e20) = e23) | ~ (op2(e20, e20) = e23)) & (~ (op2(e20, e20) = e22) | ~ (op2(e20, e20) = e22)) & (~ (op2(e20, e20) = e21) | ~ (op2(e20, e20) = e21)) & (~ (e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20)))) | ~ (e20 = op2(e20, e20))) & (e20 = op2(e20, e20))) | ~ sP75), inference(nnf_transformation, [], [f99])).
fof(f2229, plain, (~ spl102_209 | ~ spl102_126 | ~ spl102_124 | ~ spl102_112), inference(avatar_split_clause, [], [f743, f1528, f1579, f1588, f2226])).
fof(f743, plain, (~ (e23 = op2(e21, e20)) | ~ (e23 = op2(e20, e21)) | ~ (op2(e20, e20) = e21) | ~ sP76), inference(cnf_transformation, [], [f186])).
fof(f186, plain, (((((~ (e23 = op2(e21, e20)) | ~ (e23 = op2(e20, e21))) & (~ (e22 = op2(e21, e20)) | ~ (e22 = op2(e20, e21))) & (~ (e21 = op2(e21, e20)) | ~ (e21 = op2(e20, e21))) & (~ (e20 = op2(e21, e20)) | ~ (e20 = op2(e20, e21)))) | ~ (op2(e20, e20) = e21)) & (op2(e20, e20) = e21)) | ~ sP76), inference(nnf_transformation, [], [f100])).
fof(f2224, plain, (~ spl102_208 | spl102_127), inference(avatar_split_clause, [], [f734, f1592, f2217])).
fof(f1592, plain, (spl102_127 <=> (op2(e20, e20) = e22)), introduced(avatar_definition, [new_symbols(naming, [spl102_127])])).
fof(f734, plain, ((op2(e20, e20) = e22) | ~ sP77), inference(cnf_transformation, [], [f185])).
fof(f185, plain, (((((~ (e23 = op2(e22, e20)) | ~ (e23 = op2(e20, e22))) & (~ (e22 = op2(e22, e20)) | ~ (e22 = op2(e20, e22))) & (~ (e21 = op2(e22, e20)) | ~ (e21 = op2(e20, e22))) & (~ (e20 = op2(e22, e20)) | ~ (e20 = op2(e20, e22)))) | ~ (op2(e20, e20) = e22)) & (op2(e20, e20) = e22)) | ~ sP77), inference(nnf_transformation, [], [f101])).
fof(f2223, plain, (~ spl102_208 | ~ spl102_127 | ~ spl102_117 | ~ spl102_93), inference(avatar_split_clause, [], [f735, f1448, f1550, f1592, f2217])).
fof(f735, plain, (~ (e20 = op2(e22, e20)) | ~ (e20 = op2(e20, e22)) | ~ (op2(e20, e20) = e22) | ~ sP77), inference(cnf_transformation, [], [f185])).
fof(f2215, plain, (~ spl102_207 | spl102_128), inference(avatar_split_clause, [], [f729, f1596, f2208])).
fof(f1596, plain, (spl102_128 <=> (op2(e20, e20) = e23)), introduced(avatar_definition, [new_symbols(naming, [spl102_128])])).
fof(f729, plain, ((op2(e20, e20) = e23) | ~ sP78), inference(cnf_transformation, [], [f184])).
fof(f184, plain, (((((~ (e23 = op2(e23, e20)) | ~ (e23 = op2(e20, e23))) & (~ (e22 = op2(e23, e20)) | ~ (e22 = op2(e20, e23))) & (~ (e21 = op2(e23, e20)) | ~ (e21 = op2(e20, e23))) & (~ (e20 = op2(e23, e20)) | ~ (e20 = op2(e20, e23)))) | ~ (op2(e20, e20) = e23)) & (op2(e20, e20) = e23)) | ~ sP78), inference(nnf_transformation, [], [f102])).
fof(f2212, plain, (~ spl102_207 | ~ spl102_128 | ~ spl102_115 | ~ spl102_79), inference(avatar_split_clause, [], [f732, f1388, f1541, f1596, f2208])).
fof(f732, plain, (~ (e22 = op2(e23, e20)) | ~ (e22 = op2(e20, e23)) | ~ (op2(e20, e20) = e23) | ~ sP78), inference(cnf_transformation, [], [f184])).
fof(f2206, plain, (~ spl102_206 | spl102_105), inference(avatar_split_clause, [], [f724, f1499, f2199])).
fof(f724, plain, ((e20 = op2(e21, e21)) | ~ sP79), inference(cnf_transformation, [], [f183])).
fof(f183, plain, (((((~ (e23 = op2(e20, e21)) | ~ (e23 = op2(e21, e20))) & (~ (e22 = op2(e20, e21)) | ~ (e22 = op2(e21, e20))) & (~ (e21 = op2(e20, e21)) | ~ (e21 = op2(e21, e20))) & (~ (e20 = op2(e20, e21)) | ~ (e20 = op2(e21, e20)))) | ~ (e20 = op2(e21, e21))) & (e20 = op2(e21, e21))) | ~ sP79), inference(nnf_transformation, [], [f103])).
fof(f2202, plain, (~ spl102_206 | ~ spl102_105 | ~ spl102_112 | ~ spl102_124), inference(avatar_split_clause, [], [f728, f1579, f1528, f1499, f2199])).
fof(f728, plain, (~ (e23 = op2(e20, e21)) | ~ (e23 = op2(e21, e20)) | ~ (e20 = op2(e21, e21)) | ~ sP79), inference(cnf_transformation, [], [f183])).
fof(f2197, plain, (~ spl102_205 | spl102_106), inference(avatar_split_clause, [], [f719, f1503, f2190])).
fof(f719, plain, ((e21 = op2(e21, e21)) | ~ sP80), inference(cnf_transformation, [], [f182])).
fof(f182, plain, (((((~ (e23 = op2(e21, e21)) | ~ (e23 = op2(e21, e21))) & (~ (e22 = op2(e21, e21)) | ~ (e22 = op2(e21, e21))) & (~ (e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & (~ (e20 = op2(e21, e21)) | ~ (e20 = op2(e21, e21)))) | ~ (e21 = op2(e21, e21))) & (e21 = op2(e21, e21))) | ~ sP80), inference(nnf_transformation, [], [f104])).
fof(f2195, plain, (~ spl102_205 | ~ spl102_106), inference(avatar_split_clause, [], [f997, f1503, f2190])).
fof(f997, plain, (~ (e21 = op2(e21, e21)) | ~ sP80), inference(duplicate_literal_removal, [], [f721])).
fof(f721, plain, (~ (e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21)) | ~ sP80), inference(cnf_transformation, [], [f182])).
fof(f2188, plain, (~ spl102_204 | spl102_107), inference(avatar_split_clause, [], [f714, f1507, f2181])).
fof(f714, plain, ((e22 = op2(e21, e21)) | ~ sP81), inference(cnf_transformation, [], [f181])).
fof(f181, plain, (((((~ (e23 = op2(e22, e21)) | ~ (e23 = op2(e21, e22))) & (~ (e22 = op2(e22, e21)) | ~ (e22 = op2(e21, e22))) & (~ (e21 = op2(e22, e21)) | ~ (e21 = op2(e21, e22))) & (~ (e20 = op2(e22, e21)) | ~ (e20 = op2(e21, e22)))) | ~ (e22 = op2(e21, e21))) & (e22 = op2(e21, e21))) | ~ sP81), inference(nnf_transformation, [], [f105])).
fof(f2186, plain, (~ spl102_204 | ~ spl102_107 | ~ spl102_102 | ~ spl102_90), inference(avatar_split_clause, [], [f716, f1435, f1486, f1507, f2181])).
fof(f716, plain, (~ (e21 = op2(e22, e21)) | ~ (e21 = op2(e21, e22)) | ~ (e22 = op2(e21, e21)) | ~ sP81), inference(cnf_transformation, [], [f181])).
fof(f2179, plain, (~ spl102_203 | spl102_108), inference(avatar_split_clause, [], [f709, f1511, f2172])).
fof(f709, plain, ((e23 = op2(e21, e21)) | ~ sP82), inference(cnf_transformation, [], [f180])).
fof(f180, plain, (((((~ (e23 = op2(e23, e21)) | ~ (e23 = op2(e21, e23))) & (~ (e22 = op2(e23, e21)) | ~ (e22 = op2(e21, e23))) & (~ (e21 = op2(e23, e21)) | ~ (e21 = op2(e21, e23))) & (~ (e20 = op2(e23, e21)) | ~ (e20 = op2(e21, e23)))) | ~ (e23 = op2(e21, e21))) & (e23 = op2(e21, e21))) | ~ sP82), inference(nnf_transformation, [], [f106])).
fof(f2170, plain, (~ spl102_202 | spl102_85), inference(avatar_split_clause, [], [f704, f1414, f2163])).
fof(f704, plain, ((e20 = op2(e22, e22)) | ~ sP83), inference(cnf_transformation, [], [f179])).
fof(f179, plain, (((((~ (e23 = op2(e20, e22)) | ~ (e23 = op2(e22, e20))) & (~ (e22 = op2(e20, e22)) | ~ (e22 = op2(e22, e20))) & (~ (e21 = op2(e20, e22)) | ~ (e21 = op2(e22, e20))) & (~ (e20 = op2(e20, e22)) | ~ (e20 = op2(e22, e20)))) | ~ (e20 = op2(e22, e22))) & (e20 = op2(e22, e22))) | ~ sP83), inference(nnf_transformation, [], [f107])).
fof(f2161, plain, (~ spl102_201 | spl102_86), inference(avatar_split_clause, [], [f699, f1418, f2154])).
fof(f699, plain, ((e21 = op2(e22, e22)) | ~ sP84), inference(cnf_transformation, [], [f178])).
fof(f178, plain, (((((~ (e23 = op2(e21, e22)) | ~ (e23 = op2(e22, e21))) & (~ (e22 = op2(e21, e22)) | ~ (e22 = op2(e22, e21))) & (~ (e21 = op2(e21, e22)) | ~ (e21 = op2(e22, e21))) & (~ (e20 = op2(e21, e22)) | ~ (e20 = op2(e22, e21)))) | ~ (e21 = op2(e22, e22))) & (e21 = op2(e22, e22))) | ~ sP84), inference(nnf_transformation, [], [f108])).
fof(f2158, plain, (~ spl102_201 | ~ spl102_86 | ~ spl102_91 | ~ spl102_103), inference(avatar_split_clause, [], [f702, f1490, f1439, f1418, f2154])).
fof(f702, plain, (~ (e22 = op2(e21, e22)) | ~ (e22 = op2(e22, e21)) | ~ (e21 = op2(e22, e22)) | ~ sP84), inference(cnf_transformation, [], [f178])).
fof(f2152, plain, (~ spl102_200 | spl102_87), inference(avatar_split_clause, [], [f694, f1422, f2145])).
fof(f694, plain, ((e22 = op2(e22, e22)) | ~ sP85), inference(cnf_transformation, [], [f177])).
fof(f177, plain, (((((~ (e23 = op2(e22, e22)) | ~ (e23 = op2(e22, e22))) & (~ (e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & (~ (e21 = op2(e22, e22)) | ~ (e21 = op2(e22, e22))) & (~ (e20 = op2(e22, e22)) | ~ (e20 = op2(e22, e22)))) | ~ (e22 = op2(e22, e22))) & (e22 = op2(e22, e22))) | ~ sP85), inference(nnf_transformation, [], [f109])).
fof(f2149, plain, (~ spl102_200 | ~ spl102_87), inference(avatar_split_clause, [], [f1002, f1422, f2145])).
fof(f1002, plain, (~ (e22 = op2(e22, e22)) | ~ sP85), inference(duplicate_literal_removal, [], [f697])).
fof(f697, plain, (~ (e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22)) | ~ sP85), inference(cnf_transformation, [], [f177])).
fof(f2143, plain, (~ spl102_199 | spl102_88), inference(avatar_split_clause, [], [f689, f1426, f2136])).
fof(f689, plain, ((e23 = op2(e22, e22)) | ~ sP86), inference(cnf_transformation, [], [f176])).
fof(f176, plain, (((((~ (e23 = op2(e23, e22)) | ~ (e23 = op2(e22, e23))) & (~ (e22 = op2(e23, e22)) | ~ (e22 = op2(e22, e23))) & (~ (e21 = op2(e23, e22)) | ~ (e21 = op2(e22, e23))) & (~ (e20 = op2(e23, e22)) | ~ (e20 = op2(e22, e23)))) | ~ (e23 = op2(e22, e22))) & (e23 = op2(e22, e22))) | ~ sP86), inference(nnf_transformation, [], [f110])).
fof(f2141, plain, (~ spl102_199 | ~ spl102_88 | ~ spl102_82 | ~ spl102_70), inference(avatar_split_clause, [], [f691, f1350, f1401, f1426, f2136])).
fof(f691, plain, (~ (e21 = op2(e23, e22)) | ~ (e21 = op2(e22, e23)) | ~ (e23 = op2(e22, e22)) | ~ sP86), inference(cnf_transformation, [], [f176])).
fof(f2134, plain, (~ spl102_198 | spl102_65), inference(avatar_split_clause, [], [f684, f1329, f2127])).
fof(f684, plain, ((e20 = op2(e23, e23)) | ~ sP87), inference(cnf_transformation, [], [f175])).
fof(f175, plain, (((((~ (e23 = op2(e20, e23)) | ~ (e23 = op2(e23, e20))) & (~ (e22 = op2(e20, e23)) | ~ (e22 = op2(e23, e20))) & (~ (e21 = op2(e20, e23)) | ~ (e21 = op2(e23, e20))) & (~ (e20 = op2(e20, e23)) | ~ (e20 = op2(e23, e20)))) | ~ (e20 = op2(e23, e23))) & (e20 = op2(e23, e23))) | ~ sP87), inference(nnf_transformation, [], [f111])).
fof(f2131, plain, (~ spl102_198 | ~ spl102_65 | ~ spl102_79 | ~ spl102_115), inference(avatar_split_clause, [], [f687, f1541, f1388, f1329, f2127])).
fof(f687, plain, (~ (e22 = op2(e20, e23)) | ~ (e22 = op2(e23, e20)) | ~ (e20 = op2(e23, e23)) | ~ sP87), inference(cnf_transformation, [], [f175])).
fof(f2125, plain, (~ spl102_197 | spl102_66), inference(avatar_split_clause, [], [f679, f1333, f2118])).
fof(f679, plain, ((e21 = op2(e23, e23)) | ~ sP88), inference(cnf_transformation, [], [f174])).
fof(f174, plain, (((((~ (e23 = op2(e21, e23)) | ~ (e23 = op2(e23, e21))) & (~ (e22 = op2(e21, e23)) | ~ (e22 = op2(e23, e21))) & (~ (e21 = op2(e21, e23)) | ~ (e21 = op2(e23, e21))) & (~ (e20 = op2(e21, e23)) | ~ (e20 = op2(e23, e21)))) | ~ (e21 = op2(e23, e23))) & (e21 = op2(e23, e23))) | ~ sP88), inference(nnf_transformation, [], [f112])).
fof(f2124, plain, (~ spl102_197 | ~ spl102_66 | ~ spl102_73 | ~ spl102_97), inference(avatar_split_clause, [], [f680, f1465, f1363, f1333, f2118])).
fof(f680, plain, (~ (e20 = op2(e21, e23)) | ~ (e20 = op2(e23, e21)) | ~ (e21 = op2(e23, e23)) | ~ sP88), inference(cnf_transformation, [], [f174])).
fof(f2116, plain, (~ spl102_196 | spl102_67), inference(avatar_split_clause, [], [f674, f1337, f2109])).
fof(f674, plain, ((e22 = op2(e23, e23)) | ~ sP89), inference(cnf_transformation, [], [f173])).
fof(f173, plain, (((((~ (e23 = op2(e22, e23)) | ~ (e23 = op2(e23, e22))) & (~ (e22 = op2(e22, e23)) | ~ (e22 = op2(e23, e22))) & (~ (e21 = op2(e22, e23)) | ~ (e21 = op2(e23, e22))) & (~ (e20 = op2(e22, e23)) | ~ (e20 = op2(e23, e22)))) | ~ (e22 = op2(e23, e23))) & (e22 = op2(e23, e23))) | ~ sP89), inference(nnf_transformation, [], [f113])).
fof(f2107, plain, (spl102_194 | spl102_193 | spl102_192 | spl102_191 | spl102_190 | spl102_189 | spl102_188 | spl102_187 | spl102_186 | spl102_185 | spl102_184 | spl102_183 | spl102_182 | spl102_181 | spl102_180 | spl102_4), inference(avatar_split_clause, [], [f660, f1037, f1942, f1952, f1962, f1972, f1982, f1992, f2002, f2012, f2022, f2032, f2042, f2052, f2062, f2072, f2082])).
fof(f2082, plain, (spl102_194 <=> sP0), introduced(avatar_definition, [new_symbols(naming, [spl102_194])])).
fof(f2072, plain, (spl102_193 <=> sP1), introduced(avatar_definition, [new_symbols(naming, [spl102_193])])).
fof(f2062, plain, (spl102_192 <=> sP2), introduced(avatar_definition, [new_symbols(naming, [spl102_192])])).
fof(f2052, plain, (spl102_191 <=> sP3), introduced(avatar_definition, [new_symbols(naming, [spl102_191])])).
fof(f2042, plain, (spl102_190 <=> sP4), introduced(avatar_definition, [new_symbols(naming, [spl102_190])])).
fof(f2032, plain, (spl102_189 <=> sP5), introduced(avatar_definition, [new_symbols(naming, [spl102_189])])).
fof(f2022, plain, (spl102_188 <=> sP6), introduced(avatar_definition, [new_symbols(naming, [spl102_188])])).
fof(f2012, plain, (spl102_187 <=> sP7), introduced(avatar_definition, [new_symbols(naming, [spl102_187])])).
fof(f2002, plain, (spl102_186 <=> sP8), introduced(avatar_definition, [new_symbols(naming, [spl102_186])])).
fof(f1992, plain, (spl102_185 <=> sP9), introduced(avatar_definition, [new_symbols(naming, [spl102_185])])).
fof(f1982, plain, (spl102_184 <=> sP10), introduced(avatar_definition, [new_symbols(naming, [spl102_184])])).
fof(f1972, plain, (spl102_183 <=> sP11), introduced(avatar_definition, [new_symbols(naming, [spl102_183])])).
fof(f1962, plain, (spl102_182 <=> sP12), introduced(avatar_definition, [new_symbols(naming, [spl102_182])])).
fof(f1952, plain, (spl102_181 <=> sP13), introduced(avatar_definition, [new_symbols(naming, [spl102_181])])).
fof(f1942, plain, (spl102_180 <=> sP14), introduced(avatar_definition, [new_symbols(naming, [spl102_180])])).
fof(f660, plain, ((e13 = op1(e13, e13)) | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f68])).
fof(f68, plain, ((((((~ (e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13))) & (~ (e12 = op1(e13, e13)) | ~ (e12 = op1(e13, e13))) & (~ (e11 = op1(e13, e13)) | ~ (e11 = op1(e13, e13))) & (~ (e10 = op1(e13, e13)) | ~ (e10 = op1(e13, e13)))) | ~ (e13 = op1(e13, e13))) & (e13 = op1(e13, e13))) | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30) & (((e13 = op1(e13, op1(e13, e13))) & ~ (op1(e13, e13) = op1(e13, e13)) & (op1(e13, e13) = op1(e13, e13))) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15) & ((~ (e13 = op1(e13, e13)) & ~ (e12 = op1(e12, e13)) & ~ (e11 = op1(e11, e13)) & ~ (e10 = op1(e10, e13)) & (e13 = op1(e13, e13)) & (e13 = op1(e13, e13))) | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0)), inference(definition_folding, [], [f20, e67, e66, e65, e64, e63, e62, e61, e60, e59, e58, e57, e56, e55, e54, e53, e52, e51, e50, e49, e48, e47, e46, e45, e44, e43, e42, e41, e40, e39, e38, e37, e36, e35, e34, e33, e32, e31, e30, e29, e28, e27, e26, e25, e24, e23])).
fof(f23, plain, ((~ (e13 = op1(e13, e10)) & ~ (e12 = op1(e12, e10)) & ~ (e11 = op1(e11, e10)) & ~ (e10 = op1(e10, e10)) & (e10 = op1(e10, e10)) & (e10 = op1(e10, e10))) | ~ sP0), inference(usedef, [], [e23])).
fof(e23, plain, (sP0 <=> (~ (e13 = op1(e13, e10)) & ~ (e12 = op1(e12, e10)) & ~ (e11 = op1(e11, e10)) & ~ (e10 = op1(e10, e10)) & (e10 = op1(e10, e10)) & (e10 = op1(e10, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f24, plain, ((~ (e13 = op1(e13, e10)) & ~ (e12 = op1(e12, e10)) & ~ (e11 = op1(e11, e10)) & ~ (e10 = op1(e10, e10)) & (e10 = op1(e11, e10)) & (e10 = op1(e10, e11))) | ~ sP1), inference(usedef, [], [e24])).
fof(e24, plain, (sP1 <=> (~ (e13 = op1(e13, e10)) & ~ (e12 = op1(e12, e10)) & ~ (e11 = op1(e11, e10)) & ~ (e10 = op1(e10, e10)) & (e10 = op1(e11, e10)) & (e10 = op1(e10, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f25, plain, ((~ (e13 = op1(e13, e10)) & ~ (e12 = op1(e12, e10)) & ~ (e11 = op1(e11, e10)) & ~ (e10 = op1(e10, e10)) & (e10 = op1(e12, e10)) & (e10 = op1(e10, e12))) | ~ sP2), inference(usedef, [], [e25])).
fof(e25, plain, (sP2 <=> (~ (e13 = op1(e13, e10)) & ~ (e12 = op1(e12, e10)) & ~ (e11 = op1(e11, e10)) & ~ (e10 = op1(e10, e10)) & (e10 = op1(e12, e10)) & (e10 = op1(e10, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f26, plain, ((~ (e13 = op1(e13, e10)) & ~ (e12 = op1(e12, e10)) & ~ (e11 = op1(e11, e10)) & ~ (e10 = op1(e10, e10)) & (e10 = op1(e13, e10)) & (e10 = op1(e10, e13))) | ~ sP3), inference(usedef, [], [e26])).
fof(e26, plain, (sP3 <=> (~ (e13 = op1(e13, e10)) & ~ (e12 = op1(e12, e10)) & ~ (e11 = op1(e11, e10)) & ~ (e10 = op1(e10, e10)) & (e10 = op1(e13, e10)) & (e10 = op1(e10, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f27, plain, ((~ (e13 = op1(e13, e11)) & ~ (e12 = op1(e12, e11)) & ~ (e11 = op1(e11, e11)) & ~ (e10 = op1(e10, e11)) & (e11 = op1(e10, e11)) & (e11 = op1(e11, e10))) | ~ sP4), inference(usedef, [], [e27])).
fof(e27, plain, (sP4 <=> (~ (e13 = op1(e13, e11)) & ~ (e12 = op1(e12, e11)) & ~ (e11 = op1(e11, e11)) & ~ (e10 = op1(e10, e11)) & (e11 = op1(e10, e11)) & (e11 = op1(e11, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f28, plain, ((~ (e13 = op1(e13, e11)) & ~ (e12 = op1(e12, e11)) & ~ (e11 = op1(e11, e11)) & ~ (e10 = op1(e10, e11)) & (e11 = op1(e11, e11)) & (e11 = op1(e11, e11))) | ~ sP5), inference(usedef, [], [e28])).
fof(e28, plain, (sP5 <=> (~ (e13 = op1(e13, e11)) & ~ (e12 = op1(e12, e11)) & ~ (e11 = op1(e11, e11)) & ~ (e10 = op1(e10, e11)) & (e11 = op1(e11, e11)) & (e11 = op1(e11, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f29, plain, ((~ (e13 = op1(e13, e11)) & ~ (e12 = op1(e12, e11)) & ~ (e11 = op1(e11, e11)) & ~ (e10 = op1(e10, e11)) & (e11 = op1(e12, e11)) & (e11 = op1(e11, e12))) | ~ sP6), inference(usedef, [], [e29])).
fof(e29, plain, (sP6 <=> (~ (e13 = op1(e13, e11)) & ~ (e12 = op1(e12, e11)) & ~ (e11 = op1(e11, e11)) & ~ (e10 = op1(e10, e11)) & (e11 = op1(e12, e11)) & (e11 = op1(e11, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f30, plain, ((~ (e13 = op1(e13, e11)) & ~ (e12 = op1(e12, e11)) & ~ (e11 = op1(e11, e11)) & ~ (e10 = op1(e10, e11)) & (e11 = op1(e13, e11)) & (e11 = op1(e11, e13))) | ~ sP7), inference(usedef, [], [e30])).
fof(e30, plain, (sP7 <=> (~ (e13 = op1(e13, e11)) & ~ (e12 = op1(e12, e11)) & ~ (e11 = op1(e11, e11)) & ~ (e10 = op1(e10, e11)) & (e11 = op1(e13, e11)) & (e11 = op1(e11, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f31, plain, ((~ (e13 = op1(e13, e12)) & ~ (e12 = op1(e12, e12)) & ~ (e11 = op1(e11, e12)) & ~ (e10 = op1(e10, e12)) & (e12 = op1(e10, e12)) & (e12 = op1(e12, e10))) | ~ sP8), inference(usedef, [], [e31])).
fof(e31, plain, (sP8 <=> (~ (e13 = op1(e13, e12)) & ~ (e12 = op1(e12, e12)) & ~ (e11 = op1(e11, e12)) & ~ (e10 = op1(e10, e12)) & (e12 = op1(e10, e12)) & (e12 = op1(e12, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f32, plain, ((~ (e13 = op1(e13, e12)) & ~ (e12 = op1(e12, e12)) & ~ (e11 = op1(e11, e12)) & ~ (e10 = op1(e10, e12)) & (e12 = op1(e11, e12)) & (e12 = op1(e12, e11))) | ~ sP9), inference(usedef, [], [e32])).
fof(e32, plain, (sP9 <=> (~ (e13 = op1(e13, e12)) & ~ (e12 = op1(e12, e12)) & ~ (e11 = op1(e11, e12)) & ~ (e10 = op1(e10, e12)) & (e12 = op1(e11, e12)) & (e12 = op1(e12, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f33, plain, ((~ (e13 = op1(e13, e12)) & ~ (e12 = op1(e12, e12)) & ~ (e11 = op1(e11, e12)) & ~ (e10 = op1(e10, e12)) & (e12 = op1(e12, e12)) & (e12 = op1(e12, e12))) | ~ sP10), inference(usedef, [], [e33])).
fof(e33, plain, (sP10 <=> (~ (e13 = op1(e13, e12)) & ~ (e12 = op1(e12, e12)) & ~ (e11 = op1(e11, e12)) & ~ (e10 = op1(e10, e12)) & (e12 = op1(e12, e12)) & (e12 = op1(e12, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f34, plain, ((~ (e13 = op1(e13, e12)) & ~ (e12 = op1(e12, e12)) & ~ (e11 = op1(e11, e12)) & ~ (e10 = op1(e10, e12)) & (e12 = op1(e13, e12)) & (e12 = op1(e12, e13))) | ~ sP11), inference(usedef, [], [e34])).
fof(e34, plain, (sP11 <=> (~ (e13 = op1(e13, e12)) & ~ (e12 = op1(e12, e12)) & ~ (e11 = op1(e11, e12)) & ~ (e10 = op1(e10, e12)) & (e12 = op1(e13, e12)) & (e12 = op1(e12, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f35, plain, ((~ (e13 = op1(e13, e13)) & ~ (e12 = op1(e12, e13)) & ~ (e11 = op1(e11, e13)) & ~ (e10 = op1(e10, e13)) & (e13 = op1(e10, e13)) & (e13 = op1(e13, e10))) | ~ sP12), inference(usedef, [], [e35])).
fof(e35, plain, (sP12 <=> (~ (e13 = op1(e13, e13)) & ~ (e12 = op1(e12, e13)) & ~ (e11 = op1(e11, e13)) & ~ (e10 = op1(e10, e13)) & (e13 = op1(e10, e13)) & (e13 = op1(e13, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f36, plain, ((~ (e13 = op1(e13, e13)) & ~ (e12 = op1(e12, e13)) & ~ (e11 = op1(e11, e13)) & ~ (e10 = op1(e10, e13)) & (e13 = op1(e11, e13)) & (e13 = op1(e13, e11))) | ~ sP13), inference(usedef, [], [e36])).
fof(e36, plain, (sP13 <=> (~ (e13 = op1(e13, e13)) & ~ (e12 = op1(e12, e13)) & ~ (e11 = op1(e11, e13)) & ~ (e10 = op1(e10, e13)) & (e13 = op1(e11, e13)) & (e13 = op1(e13, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f37, plain, ((~ (e13 = op1(e13, e13)) & ~ (e12 = op1(e12, e13)) & ~ (e11 = op1(e11, e13)) & ~ (e10 = op1(e10, e13)) & (e13 = op1(e12, e13)) & (e13 = op1(e13, e12))) | ~ sP14), inference(usedef, [], [e37])).
fof(e37, plain, (sP14 <=> (~ (e13 = op1(e13, e13)) & ~ (e12 = op1(e12, e13)) & ~ (e11 = op1(e11, e13)) & ~ (e10 = op1(e10, e13)) & (e13 = op1(e12, e13)) & (e13 = op1(e13, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f38, plain, (((e10 = op1(e10, op1(e10, e10))) & ~ (op1(e10, e10) = op1(e10, e10)) & (op1(e10, e10) = op1(e10, e10))) | ~ sP15), inference(usedef, [], [e38])).
fof(e38, plain, (sP15 <=> ((e10 = op1(e10, op1(e10, e10))) & ~ (op1(e10, e10) = op1(e10, e10)) & (op1(e10, e10) = op1(e10, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f39, plain, (((e11 = op1(e11, op1(e10, e11))) & ~ (op1(e10, e11) = op1(e11, e11)) & (op1(e10, e11) = op1(e11, e10))) | ~ sP16), inference(usedef, [], [e39])).
fof(e39, plain, (sP16 <=> ((e11 = op1(e11, op1(e10, e11))) & ~ (op1(e10, e11) = op1(e11, e11)) & (op1(e10, e11) = op1(e11, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f40, plain, (((e12 = op1(e12, op1(e10, e12))) & ~ (op1(e10, e12) = op1(e12, e12)) & (op1(e10, e12) = op1(e12, e10))) | ~ sP17), inference(usedef, [], [e40])).
fof(e40, plain, (sP17 <=> ((e12 = op1(e12, op1(e10, e12))) & ~ (op1(e10, e12) = op1(e12, e12)) & (op1(e10, e12) = op1(e12, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f41, plain, (((e13 = op1(e13, op1(e10, e13))) & ~ (op1(e10, e13) = op1(e13, e13)) & (op1(e10, e13) = op1(e13, e10))) | ~ sP18), inference(usedef, [], [e41])).
fof(e41, plain, (sP18 <=> ((e13 = op1(e13, op1(e10, e13))) & ~ (op1(e10, e13) = op1(e13, e13)) & (op1(e10, e13) = op1(e13, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f42, plain, (((e10 = op1(e10, op1(e11, e10))) & ~ (op1(e10, e10) = op1(e11, e10)) & (op1(e10, e11) = op1(e11, e10))) | ~ sP19), inference(usedef, [], [e42])).
fof(e42, plain, (sP19 <=> ((e10 = op1(e10, op1(e11, e10))) & ~ (op1(e10, e10) = op1(e11, e10)) & (op1(e10, e11) = op1(e11, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP19])])).
fof(f43, plain, (((e11 = op1(e11, op1(e11, e11))) & ~ (op1(e11, e11) = op1(e11, e11)) & (op1(e11, e11) = op1(e11, e11))) | ~ sP20), inference(usedef, [], [e43])).
fof(e43, plain, (sP20 <=> ((e11 = op1(e11, op1(e11, e11))) & ~ (op1(e11, e11) = op1(e11, e11)) & (op1(e11, e11) = op1(e11, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP20])])).
fof(f44, plain, (((e12 = op1(e12, op1(e11, e12))) & ~ (op1(e11, e12) = op1(e12, e12)) & (op1(e11, e12) = op1(e12, e11))) | ~ sP21), inference(usedef, [], [e44])).
fof(e44, plain, (sP21 <=> ((e12 = op1(e12, op1(e11, e12))) & ~ (op1(e11, e12) = op1(e12, e12)) & (op1(e11, e12) = op1(e12, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP21])])).
fof(f45, plain, (((e13 = op1(e13, op1(e11, e13))) & ~ (op1(e11, e13) = op1(e13, e13)) & (op1(e11, e13) = op1(e13, e11))) | ~ sP22), inference(usedef, [], [e45])).
fof(e45, plain, (sP22 <=> ((e13 = op1(e13, op1(e11, e13))) & ~ (op1(e11, e13) = op1(e13, e13)) & (op1(e11, e13) = op1(e13, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP22])])).
fof(f46, plain, (((e10 = op1(e10, op1(e12, e10))) & ~ (op1(e10, e10) = op1(e12, e10)) & (op1(e10, e12) = op1(e12, e10))) | ~ sP23), inference(usedef, [], [e46])).
fof(e46, plain, (sP23 <=> ((e10 = op1(e10, op1(e12, e10))) & ~ (op1(e10, e10) = op1(e12, e10)) & (op1(e10, e12) = op1(e12, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP23])])).
fof(f47, plain, (((e11 = op1(e11, op1(e12, e11))) & ~ (op1(e11, e11) = op1(e12, e11)) & (op1(e11, e12) = op1(e12, e11))) | ~ sP24), inference(usedef, [], [e47])).
fof(e47, plain, (sP24 <=> ((e11 = op1(e11, op1(e12, e11))) & ~ (op1(e11, e11) = op1(e12, e11)) & (op1(e11, e12) = op1(e12, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP24])])).
fof(f48, plain, (((e12 = op1(e12, op1(e12, e12))) & ~ (op1(e12, e12) = op1(e12, e12)) & (op1(e12, e12) = op1(e12, e12))) | ~ sP25), inference(usedef, [], [e48])).
fof(e48, plain, (sP25 <=> ((e12 = op1(e12, op1(e12, e12))) & ~ (op1(e12, e12) = op1(e12, e12)) & (op1(e12, e12) = op1(e12, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP25])])).
fof(f49, plain, (((e13 = op1(e13, op1(e12, e13))) & ~ (op1(e12, e13) = op1(e13, e13)) & (op1(e12, e13) = op1(e13, e12))) | ~ sP26), inference(usedef, [], [e49])).
fof(e49, plain, (sP26 <=> ((e13 = op1(e13, op1(e12, e13))) & ~ (op1(e12, e13) = op1(e13, e13)) & (op1(e12, e13) = op1(e13, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP26])])).
fof(f50, plain, (((e10 = op1(e10, op1(e13, e10))) & ~ (op1(e10, e10) = op1(e13, e10)) & (op1(e10, e13) = op1(e13, e10))) | ~ sP27), inference(usedef, [], [e50])).
fof(e50, plain, (sP27 <=> ((e10 = op1(e10, op1(e13, e10))) & ~ (op1(e10, e10) = op1(e13, e10)) & (op1(e10, e13) = op1(e13, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP27])])).
fof(f51, plain, (((e11 = op1(e11, op1(e13, e11))) & ~ (op1(e11, e11) = op1(e13, e11)) & (op1(e11, e13) = op1(e13, e11))) | ~ sP28), inference(usedef, [], [e51])).
fof(e51, plain, (sP28 <=> ((e11 = op1(e11, op1(e13, e11))) & ~ (op1(e11, e11) = op1(e13, e11)) & (op1(e11, e13) = op1(e13, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP28])])).
fof(f52, plain, (((e12 = op1(e12, op1(e13, e12))) & ~ (op1(e12, e12) = op1(e13, e12)) & (op1(e12, e13) = op1(e13, e12))) | ~ sP29), inference(usedef, [], [e52])).
fof(e52, plain, (sP29 <=> ((e12 = op1(e12, op1(e13, e12))) & ~ (op1(e12, e12) = op1(e13, e12)) & (op1(e12, e13) = op1(e13, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP29])])).
fof(f53, plain, (((((~ (op1(e10, e10) = e13) | ~ (op1(e10, e10) = e13)) & (~ (op1(e10, e10) = e12) | ~ (op1(e10, e10) = e12)) & (~ (op1(e10, e10) = e11) | ~ (op1(e10, e10) = e11)) & (~ (e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10)))) | ~ (e10 = op1(e10, e10))) & (e10 = op1(e10, e10))) | ~ sP30), inference(usedef, [], [e53])).
fof(e53, plain, (sP30 <=> ((((~ (op1(e10, e10) = e13) | ~ (op1(e10, e10) = e13)) & (~ (op1(e10, e10) = e12) | ~ (op1(e10, e10) = e12)) & (~ (op1(e10, e10) = e11) | ~ (op1(e10, e10) = e11)) & (~ (e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10)))) | ~ (e10 = op1(e10, e10))) & (e10 = op1(e10, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP30])])).
fof(f54, plain, (((((~ (e13 = op1(e11, e10)) | ~ (e13 = op1(e10, e11))) & (~ (e12 = op1(e11, e10)) | ~ (e12 = op1(e10, e11))) & (~ (e11 = op1(e11, e10)) | ~ (e11 = op1(e10, e11))) & (~ (e10 = op1(e11, e10)) | ~ (e10 = op1(e10, e11)))) | ~ (op1(e10, e10) = e11)) & (op1(e10, e10) = e11)) | ~ sP31), inference(usedef, [], [e54])).
fof(e54, plain, (sP31 <=> ((((~ (e13 = op1(e11, e10)) | ~ (e13 = op1(e10, e11))) & (~ (e12 = op1(e11, e10)) | ~ (e12 = op1(e10, e11))) & (~ (e11 = op1(e11, e10)) | ~ (e11 = op1(e10, e11))) & (~ (e10 = op1(e11, e10)) | ~ (e10 = op1(e10, e11)))) | ~ (op1(e10, e10) = e11)) & (op1(e10, e10) = e11))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP31])])).
fof(f55, plain, (((((~ (e13 = op1(e12, e10)) | ~ (e13 = op1(e10, e12))) & (~ (e12 = op1(e12, e10)) | ~ (e12 = op1(e10, e12))) & (~ (e11 = op1(e12, e10)) | ~ (e11 = op1(e10, e12))) & (~ (e10 = op1(e12, e10)) | ~ (e10 = op1(e10, e12)))) | ~ (op1(e10, e10) = e12)) & (op1(e10, e10) = e12)) | ~ sP32), inference(usedef, [], [e55])).
fof(e55, plain, (sP32 <=> ((((~ (e13 = op1(e12, e10)) | ~ (e13 = op1(e10, e12))) & (~ (e12 = op1(e12, e10)) | ~ (e12 = op1(e10, e12))) & (~ (e11 = op1(e12, e10)) | ~ (e11 = op1(e10, e12))) & (~ (e10 = op1(e12, e10)) | ~ (e10 = op1(e10, e12)))) | ~ (op1(e10, e10) = e12)) & (op1(e10, e10) = e12))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP32])])).
fof(f56, plain, (((((~ (e13 = op1(e13, e10)) | ~ (e13 = op1(e10, e13))) & (~ (e12 = op1(e13, e10)) | ~ (e12 = op1(e10, e13))) & (~ (e11 = op1(e13, e10)) | ~ (e11 = op1(e10, e13))) & (~ (e10 = op1(e13, e10)) | ~ (e10 = op1(e10, e13)))) | ~ (op1(e10, e10) = e13)) & (op1(e10, e10) = e13)) | ~ sP33), inference(usedef, [], [e56])).
fof(e56, plain, (sP33 <=> ((((~ (e13 = op1(e13, e10)) | ~ (e13 = op1(e10, e13))) & (~ (e12 = op1(e13, e10)) | ~ (e12 = op1(e10, e13))) & (~ (e11 = op1(e13, e10)) | ~ (e11 = op1(e10, e13))) & (~ (e10 = op1(e13, e10)) | ~ (e10 = op1(e10, e13)))) | ~ (op1(e10, e10) = e13)) & (op1(e10, e10) = e13))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP33])])).
fof(f57, plain, (((((~ (e13 = op1(e10, e11)) | ~ (e13 = op1(e11, e10))) & (~ (e12 = op1(e10, e11)) | ~ (e12 = op1(e11, e10))) & (~ (e11 = op1(e10, e11)) | ~ (e11 = op1(e11, e10))) & (~ (e10 = op1(e10, e11)) | ~ (e10 = op1(e11, e10)))) | ~ (e10 = op1(e11, e11))) & (e10 = op1(e11, e11))) | ~ sP34), inference(usedef, [], [e57])).
fof(e57, plain, (sP34 <=> ((((~ (e13 = op1(e10, e11)) | ~ (e13 = op1(e11, e10))) & (~ (e12 = op1(e10, e11)) | ~ (e12 = op1(e11, e10))) & (~ (e11 = op1(e10, e11)) | ~ (e11 = op1(e11, e10))) & (~ (e10 = op1(e10, e11)) | ~ (e10 = op1(e11, e10)))) | ~ (e10 = op1(e11, e11))) & (e10 = op1(e11, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP34])])).
fof(f58, plain, (((((~ (e13 = op1(e11, e11)) | ~ (e13 = op1(e11, e11))) & (~ (e12 = op1(e11, e11)) | ~ (e12 = op1(e11, e11))) & (~ (e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & (~ (e10 = op1(e11, e11)) | ~ (e10 = op1(e11, e11)))) | ~ (e11 = op1(e11, e11))) & (e11 = op1(e11, e11))) | ~ sP35), inference(usedef, [], [e58])).
fof(e58, plain, (sP35 <=> ((((~ (e13 = op1(e11, e11)) | ~ (e13 = op1(e11, e11))) & (~ (e12 = op1(e11, e11)) | ~ (e12 = op1(e11, e11))) & (~ (e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & (~ (e10 = op1(e11, e11)) | ~ (e10 = op1(e11, e11)))) | ~ (e11 = op1(e11, e11))) & (e11 = op1(e11, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP35])])).
fof(f59, plain, (((((~ (e13 = op1(e12, e11)) | ~ (e13 = op1(e11, e12))) & (~ (e12 = op1(e12, e11)) | ~ (e12 = op1(e11, e12))) & (~ (e11 = op1(e12, e11)) | ~ (e11 = op1(e11, e12))) & (~ (e10 = op1(e12, e11)) | ~ (e10 = op1(e11, e12)))) | ~ (e12 = op1(e11, e11))) & (e12 = op1(e11, e11))) | ~ sP36), inference(usedef, [], [e59])).
fof(e59, plain, (sP36 <=> ((((~ (e13 = op1(e12, e11)) | ~ (e13 = op1(e11, e12))) & (~ (e12 = op1(e12, e11)) | ~ (e12 = op1(e11, e12))) & (~ (e11 = op1(e12, e11)) | ~ (e11 = op1(e11, e12))) & (~ (e10 = op1(e12, e11)) | ~ (e10 = op1(e11, e12)))) | ~ (e12 = op1(e11, e11))) & (e12 = op1(e11, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP36])])).
fof(f60, plain, (((((~ (e13 = op1(e13, e11)) | ~ (e13 = op1(e11, e13))) & (~ (e12 = op1(e13, e11)) | ~ (e12 = op1(e11, e13))) & (~ (e11 = op1(e13, e11)) | ~ (e11 = op1(e11, e13))) & (~ (e10 = op1(e13, e11)) | ~ (e10 = op1(e11, e13)))) | ~ (e13 = op1(e11, e11))) & (e13 = op1(e11, e11))) | ~ sP37), inference(usedef, [], [e60])).
fof(e60, plain, (sP37 <=> ((((~ (e13 = op1(e13, e11)) | ~ (e13 = op1(e11, e13))) & (~ (e12 = op1(e13, e11)) | ~ (e12 = op1(e11, e13))) & (~ (e11 = op1(e13, e11)) | ~ (e11 = op1(e11, e13))) & (~ (e10 = op1(e13, e11)) | ~ (e10 = op1(e11, e13)))) | ~ (e13 = op1(e11, e11))) & (e13 = op1(e11, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP37])])).
fof(f61, plain, (((((~ (e13 = op1(e10, e12)) | ~ (e13 = op1(e12, e10))) & (~ (e12 = op1(e10, e12)) | ~ (e12 = op1(e12, e10))) & (~ (e11 = op1(e10, e12)) | ~ (e11 = op1(e12, e10))) & (~ (e10 = op1(e10, e12)) | ~ (e10 = op1(e12, e10)))) | ~ (e10 = op1(e12, e12))) & (e10 = op1(e12, e12))) | ~ sP38), inference(usedef, [], [e61])).
fof(e61, plain, (sP38 <=> ((((~ (e13 = op1(e10, e12)) | ~ (e13 = op1(e12, e10))) & (~ (e12 = op1(e10, e12)) | ~ (e12 = op1(e12, e10))) & (~ (e11 = op1(e10, e12)) | ~ (e11 = op1(e12, e10))) & (~ (e10 = op1(e10, e12)) | ~ (e10 = op1(e12, e10)))) | ~ (e10 = op1(e12, e12))) & (e10 = op1(e12, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP38])])).
fof(f62, plain, (((((~ (e13 = op1(e11, e12)) | ~ (e13 = op1(e12, e11))) & (~ (e12 = op1(e11, e12)) | ~ (e12 = op1(e12, e11))) & (~ (e11 = op1(e11, e12)) | ~ (e11 = op1(e12, e11))) & (~ (e10 = op1(e11, e12)) | ~ (e10 = op1(e12, e11)))) | ~ (e11 = op1(e12, e12))) & (e11 = op1(e12, e12))) | ~ sP39), inference(usedef, [], [e62])).
fof(e62, plain, (sP39 <=> ((((~ (e13 = op1(e11, e12)) | ~ (e13 = op1(e12, e11))) & (~ (e12 = op1(e11, e12)) | ~ (e12 = op1(e12, e11))) & (~ (e11 = op1(e11, e12)) | ~ (e11 = op1(e12, e11))) & (~ (e10 = op1(e11, e12)) | ~ (e10 = op1(e12, e11)))) | ~ (e11 = op1(e12, e12))) & (e11 = op1(e12, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP39])])).
fof(f63, plain, (((((~ (e13 = op1(e12, e12)) | ~ (e13 = op1(e12, e12))) & (~ (e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & (~ (e11 = op1(e12, e12)) | ~ (e11 = op1(e12, e12))) & (~ (e10 = op1(e12, e12)) | ~ (e10 = op1(e12, e12)))) | ~ (e12 = op1(e12, e12))) & (e12 = op1(e12, e12))) | ~ sP40), inference(usedef, [], [e63])).
fof(e63, plain, (sP40 <=> ((((~ (e13 = op1(e12, e12)) | ~ (e13 = op1(e12, e12))) & (~ (e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & (~ (e11 = op1(e12, e12)) | ~ (e11 = op1(e12, e12))) & (~ (e10 = op1(e12, e12)) | ~ (e10 = op1(e12, e12)))) | ~ (e12 = op1(e12, e12))) & (e12 = op1(e12, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP40])])).
fof(f64, plain, (((((~ (e13 = op1(e13, e12)) | ~ (e13 = op1(e12, e13))) & (~ (e12 = op1(e13, e12)) | ~ (e12 = op1(e12, e13))) & (~ (e11 = op1(e13, e12)) | ~ (e11 = op1(e12, e13))) & (~ (e10 = op1(e13, e12)) | ~ (e10 = op1(e12, e13)))) | ~ (e13 = op1(e12, e12))) & (e13 = op1(e12, e12))) | ~ sP41), inference(usedef, [], [e64])).
fof(e64, plain, (sP41 <=> ((((~ (e13 = op1(e13, e12)) | ~ (e13 = op1(e12, e13))) & (~ (e12 = op1(e13, e12)) | ~ (e12 = op1(e12, e13))) & (~ (e11 = op1(e13, e12)) | ~ (e11 = op1(e12, e13))) & (~ (e10 = op1(e13, e12)) | ~ (e10 = op1(e12, e13)))) | ~ (e13 = op1(e12, e12))) & (e13 = op1(e12, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP41])])).
fof(f65, plain, (((((~ (e13 = op1(e10, e13)) | ~ (e13 = op1(e13, e10))) & (~ (e12 = op1(e10, e13)) | ~ (e12 = op1(e13, e10))) & (~ (e11 = op1(e10, e13)) | ~ (e11 = op1(e13, e10))) & (~ (e10 = op1(e10, e13)) | ~ (e10 = op1(e13, e10)))) | ~ (e10 = op1(e13, e13))) & (e10 = op1(e13, e13))) | ~ sP42), inference(usedef, [], [e65])).
fof(e65, plain, (sP42 <=> ((((~ (e13 = op1(e10, e13)) | ~ (e13 = op1(e13, e10))) & (~ (e12 = op1(e10, e13)) | ~ (e12 = op1(e13, e10))) & (~ (e11 = op1(e10, e13)) | ~ (e11 = op1(e13, e10))) & (~ (e10 = op1(e10, e13)) | ~ (e10 = op1(e13, e10)))) | ~ (e10 = op1(e13, e13))) & (e10 = op1(e13, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP42])])).
fof(f66, plain, (((((~ (e13 = op1(e11, e13)) | ~ (e13 = op1(e13, e11))) & (~ (e12 = op1(e11, e13)) | ~ (e12 = op1(e13, e11))) & (~ (e11 = op1(e11, e13)) | ~ (e11 = op1(e13, e11))) & (~ (e10 = op1(e11, e13)) | ~ (e10 = op1(e13, e11)))) | ~ (e11 = op1(e13, e13))) & (e11 = op1(e13, e13))) | ~ sP43), inference(usedef, [], [e66])).
fof(e66, plain, (sP43 <=> ((((~ (e13 = op1(e11, e13)) | ~ (e13 = op1(e13, e11))) & (~ (e12 = op1(e11, e13)) | ~ (e12 = op1(e13, e11))) & (~ (e11 = op1(e11, e13)) | ~ (e11 = op1(e13, e11))) & (~ (e10 = op1(e11, e13)) | ~ (e10 = op1(e13, e11)))) | ~ (e11 = op1(e13, e13))) & (e11 = op1(e13, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP43])])).
fof(f67, plain, (((((~ (e13 = op1(e12, e13)) | ~ (e13 = op1(e13, e12))) & (~ (e12 = op1(e12, e13)) | ~ (e12 = op1(e13, e12))) & (~ (e11 = op1(e12, e13)) | ~ (e11 = op1(e13, e12))) & (~ (e10 = op1(e12, e13)) | ~ (e10 = op1(e13, e12)))) | ~ (e12 = op1(e13, e13))) & (e12 = op1(e13, e13))) | ~ sP44), inference(usedef, [], [e67])).
fof(e67, plain, (sP44 <=> ((((~ (e13 = op1(e12, e13)) | ~ (e13 = op1(e13, e12))) & (~ (e12 = op1(e12, e13)) | ~ (e12 = op1(e13, e12))) & (~ (e11 = op1(e12, e13)) | ~ (e11 = op1(e13, e12))) & (~ (e10 = op1(e12, e13)) | ~ (e10 = op1(e13, e12)))) | ~ (e12 = op1(e13, e13))) & (e12 = op1(e13, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP44])])).
fof(f20, plain, ((((((~ (e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13))) & (~ (e12 = op1(e13, e13)) | ~ (e12 = op1(e13, e13))) & (~ (e11 = op1(e13, e13)) | ~ (e11 = op1(e13, e13))) & (~ (e10 = op1(e13, e13)) | ~ (e10 = op1(e13, e13)))) | ~ (e13 = op1(e13, e13))) & (e13 = op1(e13, e13))) | ((((~ (e13 = op1(e12, e13)) | ~ (e13 = op1(e13, e12))) & (~ (e12 = op1(e12, e13)) | ~ (e12 = op1(e13, e12))) & (~ (e11 = op1(e12, e13)) | ~ (e11 = op1(e13, e12))) & (~ (e10 = op1(e12, e13)) | ~ (e10 = op1(e13, e12)))) | ~ (e12 = op1(e13, e13))) & (e12 = op1(e13, e13))) | ((((~ (e13 = op1(e11, e13)) | ~ (e13 = op1(e13, e11))) & (~ (e12 = op1(e11, e13)) | ~ (e12 = op1(e13, e11))) & (~ (e11 = op1(e11, e13)) | ~ (e11 = op1(e13, e11))) & (~ (e10 = op1(e11, e13)) | ~ (e10 = op1(e13, e11)))) | ~ (e11 = op1(e13, e13))) & (e11 = op1(e13, e13))) | ((((~ (e13 = op1(e10, e13)) | ~ (e13 = op1(e13, e10))) & (~ (e12 = op1(e10, e13)) | ~ (e12 = op1(e13, e10))) & (~ (e11 = op1(e10, e13)) | ~ (e11 = op1(e13, e10))) & (~ (e10 = op1(e10, e13)) | ~ (e10 = op1(e13, e10)))) | ~ (e10 = op1(e13, e13))) & (e10 = op1(e13, e13))) | ((((~ (e13 = op1(e13, e12)) | ~ (e13 = op1(e12, e13))) & (~ (e12 = op1(e13, e12)) | ~ (e12 = op1(e12, e13))) & (~ (e11 = op1(e13, e12)) | ~ (e11 = op1(e12, e13))) & (~ (e10 = op1(e13, e12)) | ~ (e10 = op1(e12, e13)))) | ~ (e13 = op1(e12, e12))) & (e13 = op1(e12, e12))) | ((((~ (e13 = op1(e12, e12)) | ~ (e13 = op1(e12, e12))) & (~ (e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & (~ (e11 = op1(e12, e12)) | ~ (e11 = op1(e12, e12))) & (~ (e10 = op1(e12, e12)) | ~ (e10 = op1(e12, e12)))) | ~ (e12 = op1(e12, e12))) & (e12 = op1(e12, e12))) | ((((~ (e13 = op1(e11, e12)) | ~ (e13 = op1(e12, e11))) & (~ (e12 = op1(e11, e12)) | ~ (e12 = op1(e12, e11))) & (~ (e11 = op1(e11, e12)) | ~ (e11 = op1(e12, e11))) & (~ (e10 = op1(e11, e12)) | ~ (e10 = op1(e12, e11)))) | ~ (e11 = op1(e12, e12))) & (e11 = op1(e12, e12))) | ((((~ (e13 = op1(e10, e12)) | ~ (e13 = op1(e12, e10))) & (~ (e12 = op1(e10, e12)) | ~ (e12 = op1(e12, e10))) & (~ (e11 = op1(e10, e12)) | ~ (e11 = op1(e12, e10))) & (~ (e10 = op1(e10, e12)) | ~ (e10 = op1(e12, e10)))) | ~ (e10 = op1(e12, e12))) & (e10 = op1(e12, e12))) | ((((~ (e13 = op1(e13, e11)) | ~ (e13 = op1(e11, e13))) & (~ (e12 = op1(e13, e11)) | ~ (e12 = op1(e11, e13))) & (~ (e11 = op1(e13, e11)) | ~ (e11 = op1(e11, e13))) & (~ (e10 = op1(e13, e11)) | ~ (e10 = op1(e11, e13)))) | ~ (e13 = op1(e11, e11))) & (e13 = op1(e11, e11))) | ((((~ (e13 = op1(e12, e11)) | ~ (e13 = op1(e11, e12))) & (~ (e12 = op1(e12, e11)) | ~ (e12 = op1(e11, e12))) & (~ (e11 = op1(e12, e11)) | ~ (e11 = op1(e11, e12))) & (~ (e10 = op1(e12, e11)) | ~ (e10 = op1(e11, e12)))) | ~ (e12 = op1(e11, e11))) & (e12 = op1(e11, e11))) | ((((~ (e13 = op1(e11, e11)) | ~ (e13 = op1(e11, e11))) & (~ (e12 = op1(e11, e11)) | ~ (e12 = op1(e11, e11))) & (~ (e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & (~ (e10 = op1(e11, e11)) | ~ (e10 = op1(e11, e11)))) | ~ (e11 = op1(e11, e11))) & (e11 = op1(e11, e11))) | ((((~ (e13 = op1(e10, e11)) | ~ (e13 = op1(e11, e10))) & (~ (e12 = op1(e10, e11)) | ~ (e12 = op1(e11, e10))) & (~ (e11 = op1(e10, e11)) | ~ (e11 = op1(e11, e10))) & (~ (e10 = op1(e10, e11)) | ~ (e10 = op1(e11, e10)))) | ~ (e10 = op1(e11, e11))) & (e10 = op1(e11, e11))) | ((((~ (e13 = op1(e13, e10)) | ~ (e13 = op1(e10, e13))) & (~ (e12 = op1(e13, e10)) | ~ (e12 = op1(e10, e13))) & (~ (e11 = op1(e13, e10)) | ~ (e11 = op1(e10, e13))) & (~ (e10 = op1(e13, e10)) | ~ (e10 = op1(e10, e13)))) | ~ (op1(e10, e10) = e13)) & (op1(e10, e10) = e13)) | ((((~ (e13 = op1(e12, e10)) | ~ (e13 = op1(e10, e12))) & (~ (e12 = op1(e12, e10)) | ~ (e12 = op1(e10, e12))) & (~ (e11 = op1(e12, e10)) | ~ (e11 = op1(e10, e12))) & (~ (e10 = op1(e12, e10)) | ~ (e10 = op1(e10, e12)))) | ~ (op1(e10, e10) = e12)) & (op1(e10, e10) = e12)) | ((((~ (e13 = op1(e11, e10)) | ~ (e13 = op1(e10, e11))) & (~ (e12 = op1(e11, e10)) | ~ (e12 = op1(e10, e11))) & (~ (e11 = op1(e11, e10)) | ~ (e11 = op1(e10, e11))) & (~ (e10 = op1(e11, e10)) | ~ (e10 = op1(e10, e11)))) | ~ (op1(e10, e10) = e11)) & (op1(e10, e10) = e11)) | ((((~ (op1(e10, e10) = e13) | ~ (op1(e10, e10) = e13)) & (~ (op1(e10, e10) = e12) | ~ (op1(e10, e10) = e12)) & (~ (op1(e10, e10) = e11) | ~ (op1(e10, e10) = e11)) & (~ (e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10)))) | ~ (e10 = op1(e10, e10))) & (e10 = op1(e10, e10)))) & (((e13 = op1(e13, op1(e13, e13))) & ~ (op1(e13, e13) = op1(e13, e13)) & (op1(e13, e13) = op1(e13, e13))) | ((e12 = op1(e12, op1(e13, e12))) & ~ (op1(e12, e12) = op1(e13, e12)) & (op1(e12, e13) = op1(e13, e12))) | ((e11 = op1(e11, op1(e13, e11))) & ~ (op1(e11, e11) = op1(e13, e11)) & (op1(e11, e13) = op1(e13, e11))) | ((e10 = op1(e10, op1(e13, e10))) & ~ (op1(e10, e10) = op1(e13, e10)) & (op1(e10, e13) = op1(e13, e10))) | ((e13 = op1(e13, op1(e12, e13))) & ~ (op1(e12, e13) = op1(e13, e13)) & (op1(e12, e13) = op1(e13, e12))) | ((e12 = op1(e12, op1(e12, e12))) & ~ (op1(e12, e12) = op1(e12, e12)) & (op1(e12, e12) = op1(e12, e12))) | ((e11 = op1(e11, op1(e12, e11))) & ~ (op1(e11, e11) = op1(e12, e11)) & (op1(e11, e12) = op1(e12, e11))) | ((e10 = op1(e10, op1(e12, e10))) & ~ (op1(e10, e10) = op1(e12, e10)) & (op1(e10, e12) = op1(e12, e10))) | ((e13 = op1(e13, op1(e11, e13))) & ~ (op1(e11, e13) = op1(e13, e13)) & (op1(e11, e13) = op1(e13, e11))) | ((e12 = op1(e12, op1(e11, e12))) & ~ (op1(e11, e12) = op1(e12, e12)) & (op1(e11, e12) = op1(e12, e11))) | ((e11 = op1(e11, op1(e11, e11))) & ~ (op1(e11, e11) = op1(e11, e11)) & (op1(e11, e11) = op1(e11, e11))) | ((e10 = op1(e10, op1(e11, e10))) & ~ (op1(e10, e10) = op1(e11, e10)) & (op1(e10, e11) = op1(e11, e10))) | ((e13 = op1(e13, op1(e10, e13))) & ~ (op1(e10, e13) = op1(e13, e13)) & (op1(e10, e13) = op1(e13, e10))) | ((e12 = op1(e12, op1(e10, e12))) & ~ (op1(e10, e12) = op1(e12, e12)) & (op1(e10, e12) = op1(e12, e10))) | ((e11 = op1(e11, op1(e10, e11))) & ~ (op1(e10, e11) = op1(e11, e11)) & (op1(e10, e11) = op1(e11, e10))) | ((e10 = op1(e10, op1(e10, e10))) & ~ (op1(e10, e10) = op1(e10, e10)) & (op1(e10, e10) = op1(e10, e10)))) & ((~ (e13 = op1(e13, e13)) & ~ (e12 = op1(e12, e13)) & ~ (e11 = op1(e11, e13)) & ~ (e10 = op1(e10, e13)) & (e13 = op1(e13, e13)) & (e13 = op1(e13, e13))) | (~ (e13 = op1(e13, e13)) & ~ (e12 = op1(e12, e13)) & ~ (e11 = op1(e11, e13)) & ~ (e10 = op1(e10, e13)) & (e13 = op1(e12, e13)) & (e13 = op1(e13, e12))) | (~ (e13 = op1(e13, e13)) & ~ (e12 = op1(e12, e13)) & ~ (e11 = op1(e11, e13)) & ~ (e10 = op1(e10, e13)) & (e13 = op1(e11, e13)) & (e13 = op1(e13, e11))) | (~ (e13 = op1(e13, e13)) & ~ (e12 = op1(e12, e13)) & ~ (e11 = op1(e11, e13)) & ~ (e10 = op1(e10, e13)) & (e13 = op1(e10, e13)) & (e13 = op1(e13, e10))) | (~ (e13 = op1(e13, e12)) & ~ (e12 = op1(e12, e12)) & ~ (e11 = op1(e11, e12)) & ~ (e10 = op1(e10, e12)) & (e12 = op1(e13, e12)) & (e12 = op1(e12, e13))) | (~ (e13 = op1(e13, e12)) & ~ (e12 = op1(e12, e12)) & ~ (e11 = op1(e11, e12)) & ~ (e10 = op1(e10, e12)) & (e12 = op1(e12, e12)) & (e12 = op1(e12, e12))) | (~ (e13 = op1(e13, e12)) & ~ (e12 = op1(e12, e12)) & ~ (e11 = op1(e11, e12)) & ~ (e10 = op1(e10, e12)) & (e12 = op1(e11, e12)) & (e12 = op1(e12, e11))) | (~ (e13 = op1(e13, e12)) & ~ (e12 = op1(e12, e12)) & ~ (e11 = op1(e11, e12)) & ~ (e10 = op1(e10, e12)) & (e12 = op1(e10, e12)) & (e12 = op1(e12, e10))) | (~ (e13 = op1(e13, e11)) & ~ (e12 = op1(e12, e11)) & ~ (e11 = op1(e11, e11)) & ~ (e10 = op1(e10, e11)) & (e11 = op1(e13, e11)) & (e11 = op1(e11, e13))) | (~ (e13 = op1(e13, e11)) & ~ (e12 = op1(e12, e11)) & ~ (e11 = op1(e11, e11)) & ~ (e10 = op1(e10, e11)) & (e11 = op1(e12, e11)) & (e11 = op1(e11, e12))) | (~ (e13 = op1(e13, e11)) & ~ (e12 = op1(e12, e11)) & ~ (e11 = op1(e11, e11)) & ~ (e10 = op1(e10, e11)) & (e11 = op1(e11, e11)) & (e11 = op1(e11, e11))) | (~ (e13 = op1(e13, e11)) & ~ (e12 = op1(e12, e11)) & ~ (e11 = op1(e11, e11)) & ~ (e10 = op1(e10, e11)) & (e11 = op1(e10, e11)) & (e11 = op1(e11, e10))) | (~ (e13 = op1(e13, e10)) & ~ (e12 = op1(e12, e10)) & ~ (e11 = op1(e11, e10)) & ~ (e10 = op1(e10, e10)) & (e10 = op1(e13, e10)) & (e10 = op1(e10, e13))) | (~ (e13 = op1(e13, e10)) & ~ (e12 = op1(e12, e10)) & ~ (e11 = op1(e11, e10)) & ~ (e10 = op1(e10, e10)) & (e10 = op1(e12, e10)) & (e10 = op1(e10, e12))) | (~ (e13 = op1(e13, e10)) & ~ (e12 = op1(e12, e10)) & ~ (e11 = op1(e11, e10)) & ~ (e10 = op1(e10, e10)) & (e10 = op1(e11, e10)) & (e10 = op1(e10, e11))) | (~ (e13 = op1(e13, e10)) & ~ (e12 = op1(e12, e10)) & ~ (e11 = op1(e11, e10)) & ~ (e10 = op1(e10, e10)) & (e10 = op1(e10, e10)) & (e10 = op1(e10, e10))))), inference(ennf_transformation, [], [f10])).
fof(f10, plain, (((~ ((((e13 = op1(e13, e13)) & (e13 = op1(e13, e13))) | ((e12 = op1(e13, e13)) & (e12 = op1(e13, e13))) | ((e11 = op1(e13, e13)) & (e11 = op1(e13, e13))) | ((e10 = op1(e13, e13)) & (e10 = op1(e13, e13)))) & (e13 = op1(e13, e13))) & (e13 = op1(e13, e13))) | (~ ((((e13 = op1(e12, e13)) & (e13 = op1(e13, e12))) | ((e12 = op1(e12, e13)) & (e12 = op1(e13, e12))) | ((e11 = op1(e12, e13)) & (e11 = op1(e13, e12))) | ((e10 = op1(e12, e13)) & (e10 = op1(e13, e12)))) & (e12 = op1(e13, e13))) & (e12 = op1(e13, e13))) | (~ ((((e13 = op1(e11, e13)) & (e13 = op1(e13, e11))) | ((e12 = op1(e11, e13)) & (e12 = op1(e13, e11))) | ((e11 = op1(e11, e13)) & (e11 = op1(e13, e11))) | ((e10 = op1(e11, e13)) & (e10 = op1(e13, e11)))) & (e11 = op1(e13, e13))) & (e11 = op1(e13, e13))) | (~ ((((e13 = op1(e10, e13)) & (e13 = op1(e13, e10))) | ((e12 = op1(e10, e13)) & (e12 = op1(e13, e10))) | ((e11 = op1(e10, e13)) & (e11 = op1(e13, e10))) | ((e10 = op1(e10, e13)) & (e10 = op1(e13, e10)))) & (e10 = op1(e13, e13))) & (e10 = op1(e13, e13))) | (~ ((((e13 = op1(e13, e12)) & (e13 = op1(e12, e13))) | ((e12 = op1(e13, e12)) & (e12 = op1(e12, e13))) | ((e11 = op1(e13, e12)) & (e11 = op1(e12, e13))) | ((e10 = op1(e13, e12)) & (e10 = op1(e12, e13)))) & (e13 = op1(e12, e12))) & (e13 = op1(e12, e12))) | (~ ((((e13 = op1(e12, e12)) & (e13 = op1(e12, e12))) | ((e12 = op1(e12, e12)) & (e12 = op1(e12, e12))) | ((e11 = op1(e12, e12)) & (e11 = op1(e12, e12))) | ((e10 = op1(e12, e12)) & (e10 = op1(e12, e12)))) & (e12 = op1(e12, e12))) & (e12 = op1(e12, e12))) | (~ ((((e13 = op1(e11, e12)) & (e13 = op1(e12, e11))) | ((e12 = op1(e11, e12)) & (e12 = op1(e12, e11))) | ((e11 = op1(e11, e12)) & (e11 = op1(e12, e11))) | ((e10 = op1(e11, e12)) & (e10 = op1(e12, e11)))) & (e11 = op1(e12, e12))) & (e11 = op1(e12, e12))) | (~ ((((e13 = op1(e10, e12)) & (e13 = op1(e12, e10))) | ((e12 = op1(e10, e12)) & (e12 = op1(e12, e10))) | ((e11 = op1(e10, e12)) & (e11 = op1(e12, e10))) | ((e10 = op1(e10, e12)) & (e10 = op1(e12, e10)))) & (e10 = op1(e12, e12))) & (e10 = op1(e12, e12))) | (~ ((((e13 = op1(e13, e11)) & (e13 = op1(e11, e13))) | ((e12 = op1(e13, e11)) & (e12 = op1(e11, e13))) | ((e11 = op1(e13, e11)) & (e11 = op1(e11, e13))) | ((e10 = op1(e13, e11)) & (e10 = op1(e11, e13)))) & (e13 = op1(e11, e11))) & (e13 = op1(e11, e11))) | (~ ((((e13 = op1(e12, e11)) & (e13 = op1(e11, e12))) | ((e12 = op1(e12, e11)) & (e12 = op1(e11, e12))) | ((e11 = op1(e12, e11)) & (e11 = op1(e11, e12))) | ((e10 = op1(e12, e11)) & (e10 = op1(e11, e12)))) & (e12 = op1(e11, e11))) & (e12 = op1(e11, e11))) | (~ ((((e13 = op1(e11, e11)) & (e13 = op1(e11, e11))) | ((e12 = op1(e11, e11)) & (e12 = op1(e11, e11))) | ((e11 = op1(e11, e11)) & (e11 = op1(e11, e11))) | ((e10 = op1(e11, e11)) & (e10 = op1(e11, e11)))) & (e11 = op1(e11, e11))) & (e11 = op1(e11, e11))) | (~ ((((e13 = op1(e10, e11)) & (e13 = op1(e11, e10))) | ((e12 = op1(e10, e11)) & (e12 = op1(e11, e10))) | ((e11 = op1(e10, e11)) & (e11 = op1(e11, e10))) | ((e10 = op1(e10, e11)) & (e10 = op1(e11, e10)))) & (e10 = op1(e11, e11))) & (e10 = op1(e11, e11))) | (~ ((((e13 = op1(e13, e10)) & (e13 = op1(e10, e13))) | ((e12 = op1(e13, e10)) & (e12 = op1(e10, e13))) | ((e11 = op1(e13, e10)) & (e11 = op1(e10, e13))) | ((e10 = op1(e13, e10)) & (e10 = op1(e10, e13)))) & (op1(e10, e10) = e13)) & (op1(e10, e10) = e13)) | (~ ((((e13 = op1(e12, e10)) & (e13 = op1(e10, e12))) | ((e12 = op1(e12, e10)) & (e12 = op1(e10, e12))) | ((e11 = op1(e12, e10)) & (e11 = op1(e10, e12))) | ((e10 = op1(e12, e10)) & (e10 = op1(e10, e12)))) & (op1(e10, e10) = e12)) & (op1(e10, e10) = e12)) | (~ ((((e13 = op1(e11, e10)) & (e13 = op1(e10, e11))) | ((e12 = op1(e11, e10)) & (e12 = op1(e10, e11))) | ((e11 = op1(e11, e10)) & (e11 = op1(e10, e11))) | ((e10 = op1(e11, e10)) & (e10 = op1(e10, e11)))) & (op1(e10, e10) = e11)) & (op1(e10, e10) = e11)) | (~ ((((op1(e10, e10) = e13) & (op1(e10, e10) = e13)) | ((op1(e10, e10) = e12) & (op1(e10, e10) = e12)) | ((op1(e10, e10) = e11) & (op1(e10, e10) = e11)) | ((e10 = op1(e10, e10)) & (e10 = op1(e10, e10)))) & (e10 = op1(e10, e10))) & (e10 = op1(e10, e10)))) & (((e13 = op1(e13, op1(e13, e13))) & ~ (op1(e13, e13) = op1(e13, e13)) & (op1(e13, e13) = op1(e13, e13))) | ((e12 = op1(e12, op1(e13, e12))) & ~ (op1(e12, e12) = op1(e13, e12)) & (op1(e12, e13) = op1(e13, e12))) | ((e11 = op1(e11, op1(e13, e11))) & ~ (op1(e11, e11) = op1(e13, e11)) & (op1(e11, e13) = op1(e13, e11))) | ((e10 = op1(e10, op1(e13, e10))) & ~ (op1(e10, e10) = op1(e13, e10)) & (op1(e10, e13) = op1(e13, e10))) | ((e13 = op1(e13, op1(e12, e13))) & ~ (op1(e12, e13) = op1(e13, e13)) & (op1(e12, e13) = op1(e13, e12))) | ((e12 = op1(e12, op1(e12, e12))) & ~ (op1(e12, e12) = op1(e12, e12)) & (op1(e12, e12) = op1(e12, e12))) | ((e11 = op1(e11, op1(e12, e11))) & ~ (op1(e11, e11) = op1(e12, e11)) & (op1(e11, e12) = op1(e12, e11))) | ((e10 = op1(e10, op1(e12, e10))) & ~ (op1(e10, e10) = op1(e12, e10)) & (op1(e10, e12) = op1(e12, e10))) | ((e13 = op1(e13, op1(e11, e13))) & ~ (op1(e11, e13) = op1(e13, e13)) & (op1(e11, e13) = op1(e13, e11))) | ((e12 = op1(e12, op1(e11, e12))) & ~ (op1(e11, e12) = op1(e12, e12)) & (op1(e11, e12) = op1(e12, e11))) | ((e11 = op1(e11, op1(e11, e11))) & ~ (op1(e11, e11) = op1(e11, e11)) & (op1(e11, e11) = op1(e11, e11))) | ((e10 = op1(e10, op1(e11, e10))) & ~ (op1(e10, e10) = op1(e11, e10)) & (op1(e10, e11) = op1(e11, e10))) | ((e13 = op1(e13, op1(e10, e13))) & ~ (op1(e10, e13) = op1(e13, e13)) & (op1(e10, e13) = op1(e13, e10))) | ((e12 = op1(e12, op1(e10, e12))) & ~ (op1(e10, e12) = op1(e12, e12)) & (op1(e10, e12) = op1(e12, e10))) | ((e11 = op1(e11, op1(e10, e11))) & ~ (op1(e10, e11) = op1(e11, e11)) & (op1(e10, e11) = op1(e11, e10))) | ((e10 = op1(e10, op1(e10, e10))) & ~ (op1(e10, e10) = op1(e10, e10)) & (op1(e10, e10) = op1(e10, e10)))) & ((~ (e13 = op1(e13, e13)) & ~ (e12 = op1(e12, e13)) & ~ (e11 = op1(e11, e13)) & ~ (e10 = op1(e10, e13)) & (e13 = op1(e13, e13)) & (e13 = op1(e13, e13))) | (~ (e13 = op1(e13, e13)) & ~ (e12 = op1(e12, e13)) & ~ (e11 = op1(e11, e13)) & ~ (e10 = op1(e10, e13)) & (e13 = op1(e12, e13)) & (e13 = op1(e13, e12))) | (~ (e13 = op1(e13, e13)) & ~ (e12 = op1(e12, e13)) & ~ (e11 = op1(e11, e13)) & ~ (e10 = op1(e10, e13)) & (e13 = op1(e11, e13)) & (e13 = op1(e13, e11))) | (~ (e13 = op1(e13, e13)) & ~ (e12 = op1(e12, e13)) & ~ (e11 = op1(e11, e13)) & ~ (e10 = op1(e10, e13)) & (e13 = op1(e10, e13)) & (e13 = op1(e13, e10))) | (~ (e13 = op1(e13, e12)) & ~ (e12 = op1(e12, e12)) & ~ (e11 = op1(e11, e12)) & ~ (e10 = op1(e10, e12)) & (e12 = op1(e13, e12)) & (e12 = op1(e12, e13))) | (~ (e13 = op1(e13, e12)) & ~ (e12 = op1(e12, e12)) & ~ (e11 = op1(e11, e12)) & ~ (e10 = op1(e10, e12)) & (e12 = op1(e12, e12)) & (e12 = op1(e12, e12))) | (~ (e13 = op1(e13, e12)) & ~ (e12 = op1(e12, e12)) & ~ (e11 = op1(e11, e12)) & ~ (e10 = op1(e10, e12)) & (e12 = op1(e11, e12)) & (e12 = op1(e12, e11))) | (~ (e13 = op1(e13, e12)) & ~ (e12 = op1(e12, e12)) & ~ (e11 = op1(e11, e12)) & ~ (e10 = op1(e10, e12)) & (e12 = op1(e10, e12)) & (e12 = op1(e12, e10))) | (~ (e13 = op1(e13, e11)) & ~ (e12 = op1(e12, e11)) & ~ (e11 = op1(e11, e11)) & ~ (e10 = op1(e10, e11)) & (e11 = op1(e13, e11)) & (e11 = op1(e11, e13))) | (~ (e13 = op1(e13, e11)) & ~ (e12 = op1(e12, e11)) & ~ (e11 = op1(e11, e11)) & ~ (e10 = op1(e10, e11)) & (e11 = op1(e12, e11)) & (e11 = op1(e11, e12))) | (~ (e13 = op1(e13, e11)) & ~ (e12 = op1(e12, e11)) & ~ (e11 = op1(e11, e11)) & ~ (e10 = op1(e10, e11)) & (e11 = op1(e11, e11)) & (e11 = op1(e11, e11))) | (~ (e13 = op1(e13, e11)) & ~ (e12 = op1(e12, e11)) & ~ (e11 = op1(e11, e11)) & ~ (e10 = op1(e10, e11)) & (e11 = op1(e10, e11)) & (e11 = op1(e11, e10))) | (~ (e13 = op1(e13, e10)) & ~ (e12 = op1(e12, e10)) & ~ (e11 = op1(e11, e10)) & ~ (e10 = op1(e10, e10)) & (e10 = op1(e13, e10)) & (e10 = op1(e10, e13))) | (~ (e13 = op1(e13, e10)) & ~ (e12 = op1(e12, e10)) & ~ (e11 = op1(e11, e10)) & ~ (e10 = op1(e10, e10)) & (e10 = op1(e12, e10)) & (e10 = op1(e10, e12))) | (~ (e13 = op1(e13, e10)) & ~ (e12 = op1(e12, e10)) & ~ (e11 = op1(e11, e10)) & ~ (e10 = op1(e10, e10)) & (e10 = op1(e11, e10)) & (e10 = op1(e10, e11))) | (~ (e13 = op1(e13, e10)) & ~ (e12 = op1(e12, e10)) & ~ (e11 = op1(e11, e10)) & ~ (e10 = op1(e10, e10)) & (e10 = op1(e10, e10)) & (e10 = op1(e10, e10))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG109+1.p', ax10)).
fof(f2104, plain, (spl102_194 | spl102_193 | spl102_192 | spl102_191 | spl102_190 | spl102_189 | spl102_188 | spl102_187 | spl102_186 | spl102_185 | spl102_184 | spl102_183 | spl102_182 | spl102_181 | spl102_180 | ~ spl102_34), inference(avatar_split_clause, [], [f663, f1165, f1942, f1952, f1962, f1972, f1982, f1992, f2002, f2012, f2022, f2032, f2042, f2052, f2062, f2072, f2082])).
fof(f663, plain, (~ (e11 = op1(e11, e13)) | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f68])).
fof(f2102, plain, (spl102_194 | spl102_193 | spl102_192 | spl102_191 | spl102_190 | spl102_189 | spl102_188 | spl102_187 | spl102_186 | spl102_185 | spl102_184 | spl102_183 | spl102_182 | spl102_181 | spl102_180 | ~ spl102_4), inference(avatar_split_clause, [], [f665, f1037, f1942, f1952, f1962, f1972, f1982, f1992, f2002, f2012, f2022, f2032, f2042, f2052, f2062, f2072, f2082])).
fof(f665, plain, (~ (e13 = op1(e13, e13)) | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f68])).
fof(f2101, plain, (spl102_178 | spl102_176 | spl102_174 | spl102_172 | spl102_169 | spl102_167 | spl102_165 | spl102_163 | spl102_160 | spl102_157 | spl102_155 | spl102_153 | spl102_150 | spl102_147 | spl102_144), inference(avatar_split_clause, [], [f1004, f1768, f1782, f1796, f1810, f1820, f1830, f1844, f1858, f1868, f1878, f1888, f1902, f1912, f1922, f1932])).
fof(f1932, plain, (spl102_178 <=> sP15), introduced(avatar_definition, [new_symbols(naming, [spl102_178])])).
fof(f1922, plain, (spl102_176 <=> sP16), introduced(avatar_definition, [new_symbols(naming, [spl102_176])])).
fof(f1912, plain, (spl102_174 <=> sP17), introduced(avatar_definition, [new_symbols(naming, [spl102_174])])).
fof(f1902, plain, (spl102_172 <=> sP18), introduced(avatar_definition, [new_symbols(naming, [spl102_172])])).
fof(f1888, plain, (spl102_169 <=> sP19), introduced(avatar_definition, [new_symbols(naming, [spl102_169])])).
fof(f1878, plain, (spl102_167 <=> sP20), introduced(avatar_definition, [new_symbols(naming, [spl102_167])])).
fof(f1868, plain, (spl102_165 <=> sP21), introduced(avatar_definition, [new_symbols(naming, [spl102_165])])).
fof(f1858, plain, (spl102_163 <=> sP22), introduced(avatar_definition, [new_symbols(naming, [spl102_163])])).
fof(f1844, plain, (spl102_160 <=> sP23), introduced(avatar_definition, [new_symbols(naming, [spl102_160])])).
fof(f1830, plain, (spl102_157 <=> sP24), introduced(avatar_definition, [new_symbols(naming, [spl102_157])])).
fof(f1820, plain, (spl102_155 <=> sP25), introduced(avatar_definition, [new_symbols(naming, [spl102_155])])).
fof(f1810, plain, (spl102_153 <=> sP26), introduced(avatar_definition, [new_symbols(naming, [spl102_153])])).
fof(f1796, plain, (spl102_150 <=> sP27), introduced(avatar_definition, [new_symbols(naming, [spl102_150])])).
fof(f1782, plain, (spl102_147 <=> sP28), introduced(avatar_definition, [new_symbols(naming, [spl102_147])])).
fof(f1768, plain, (spl102_144 <=> sP29), introduced(avatar_definition, [new_symbols(naming, [spl102_144])])).
fof(f1004, plain, (sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15), inference(trivial_inequality_removal, [], [f667])).
fof(f667, plain, (~ (op1(e13, e13) = op1(e13, e13)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15), inference(cnf_transformation, [], [f68])).
fof(f2095, plain, (spl102_143 | spl102_142 | spl102_141 | spl102_140 | spl102_139 | spl102_138 | spl102_137 | spl102_136 | spl102_135 | spl102_134 | spl102_133 | spl102_132 | spl102_131 | spl102_130 | spl102_129 | spl102_4), inference(avatar_split_clause, [], [f669, f1037, f1633, f1642, f1651, f1660, f1669, f1678, f1687, f1696, f1705, f1714, f1723, f1732, f1741, f1750, f1759])).
fof(f1759, plain, (spl102_143 <=> sP30), introduced(avatar_definition, [new_symbols(naming, [spl102_143])])).
fof(f1750, plain, (spl102_142 <=> sP31), introduced(avatar_definition, [new_symbols(naming, [spl102_142])])).
fof(f1741, plain, (spl102_141 <=> sP32), introduced(avatar_definition, [new_symbols(naming, [spl102_141])])).
fof(f1732, plain, (spl102_140 <=> sP33), introduced(avatar_definition, [new_symbols(naming, [spl102_140])])).
fof(f1723, plain, (spl102_139 <=> sP34), introduced(avatar_definition, [new_symbols(naming, [spl102_139])])).
fof(f1714, plain, (spl102_138 <=> sP35), introduced(avatar_definition, [new_symbols(naming, [spl102_138])])).
fof(f1705, plain, (spl102_137 <=> sP36), introduced(avatar_definition, [new_symbols(naming, [spl102_137])])).
fof(f1696, plain, (spl102_136 <=> sP37), introduced(avatar_definition, [new_symbols(naming, [spl102_136])])).
fof(f1687, plain, (spl102_135 <=> sP38), introduced(avatar_definition, [new_symbols(naming, [spl102_135])])).
fof(f1678, plain, (spl102_134 <=> sP39), introduced(avatar_definition, [new_symbols(naming, [spl102_134])])).
fof(f1669, plain, (spl102_133 <=> sP40), introduced(avatar_definition, [new_symbols(naming, [spl102_133])])).
fof(f1660, plain, (spl102_132 <=> sP41), introduced(avatar_definition, [new_symbols(naming, [spl102_132])])).
fof(f1651, plain, (spl102_131 <=> sP42), introduced(avatar_definition, [new_symbols(naming, [spl102_131])])).
fof(f1642, plain, (spl102_130 <=> sP43), introduced(avatar_definition, [new_symbols(naming, [spl102_130])])).
fof(f1633, plain, (spl102_129 <=> sP44), introduced(avatar_definition, [new_symbols(naming, [spl102_129])])).
fof(f669, plain, ((e13 = op1(e13, e13)) | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30), inference(cnf_transformation, [], [f68])).
fof(f2091, plain, (spl102_143 | spl102_142 | spl102_141 | spl102_140 | spl102_139 | spl102_138 | spl102_137 | spl102_136 | spl102_135 | spl102_134 | spl102_133 | spl102_132 | spl102_131 | spl102_130 | spl102_129 | ~ spl102_4), inference(avatar_split_clause, [], [f1008, f1037, f1633, f1642, f1651, f1660, f1669, f1678, f1687, f1696, f1705, f1714, f1723, f1732, f1741, f1750, f1759])).
fof(f1008, plain, (~ (e13 = op1(e13, e13)) | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30), inference(duplicate_literal_removal, [], [f673])).
fof(f673, plain, (~ (e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13)) | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30), inference(cnf_transformation, [], [f68])).
fof(f2090, plain, (~ spl102_194 | spl102_61), inference(avatar_split_clause, [], [f654, f1280, f2082])).
fof(f654, plain, ((e10 = op1(e10, e10)) | ~ sP0), inference(cnf_transformation, [], [f172])).
fof(f172, plain, ((~ (e13 = op1(e13, e10)) & ~ (e12 = op1(e12, e10)) & ~ (e11 = op1(e11, e10)) & ~ (e10 = op1(e10, e10)) & (e10 = op1(e10, e10)) & (e10 = op1(e10, e10))) | ~ sP0), inference(nnf_transformation, [], [f23])).
fof(f2080, plain, (~ spl102_193 | spl102_57), inference(avatar_split_clause, [], [f648, f1263, f2072])).
fof(f648, plain, ((e10 = op1(e10, e11)) | ~ sP1), inference(cnf_transformation, [], [f171])).
fof(f171, plain, ((~ (e13 = op1(e13, e10)) & ~ (e12 = op1(e12, e10)) & ~ (e11 = op1(e11, e10)) & ~ (e10 = op1(e10, e10)) & (e10 = op1(e11, e10)) & (e10 = op1(e10, e11))) | ~ sP1), inference(nnf_transformation, [], [f24])).
fof(f2070, plain, (~ spl102_192 | spl102_53), inference(avatar_split_clause, [], [f642, f1246, f2062])).
fof(f642, plain, ((e10 = op1(e10, e12)) | ~ sP2), inference(cnf_transformation, [], [f170])).
fof(f170, plain, ((~ (e13 = op1(e13, e10)) & ~ (e12 = op1(e12, e10)) & ~ (e11 = op1(e11, e10)) & ~ (e10 = op1(e10, e10)) & (e10 = op1(e12, e10)) & (e10 = op1(e10, e12))) | ~ sP2), inference(nnf_transformation, [], [f25])).
fof(f2069, plain, (~ spl102_192 | spl102_29), inference(avatar_split_clause, [], [f643, f1144, f2062])).
fof(f643, plain, ((e10 = op1(e12, e10)) | ~ sP2), inference(cnf_transformation, [], [f170])).
fof(f2059, plain, (~ spl102_191 | spl102_13), inference(avatar_split_clause, [], [f637, f1076, f2052])).
fof(f637, plain, ((e10 = op1(e13, e10)) | ~ sP3), inference(cnf_transformation, [], [f169])).
fof(f169, plain, ((~ (e13 = op1(e13, e10)) & ~ (e12 = op1(e12, e10)) & ~ (e11 = op1(e11, e10)) & ~ (e10 = op1(e10, e10)) & (e10 = op1(e13, e10)) & (e10 = op1(e10, e13))) | ~ sP3), inference(nnf_transformation, [], [f26])).
fof(f2050, plain, (~ spl102_190 | spl102_46), inference(avatar_split_clause, [], [f630, f1216, f2042])).
fof(f630, plain, ((e11 = op1(e11, e10)) | ~ sP4), inference(cnf_transformation, [], [f168])).
fof(f168, plain, ((~ (e13 = op1(e13, e11)) & ~ (e12 = op1(e12, e11)) & ~ (e11 = op1(e11, e11)) & ~ (e10 = op1(e10, e11)) & (e11 = op1(e10, e11)) & (e11 = op1(e11, e10))) | ~ sP4), inference(nnf_transformation, [], [f27])).
fof(f2040, plain, (~ spl102_189 | spl102_42), inference(avatar_split_clause, [], [f624, f1199, f2032])).
fof(f624, plain, ((e11 = op1(e11, e11)) | ~ sP5), inference(cnf_transformation, [], [f167])).
fof(f167, plain, ((~ (e13 = op1(e13, e11)) & ~ (e12 = op1(e12, e11)) & ~ (e11 = op1(e11, e11)) & ~ (e10 = op1(e10, e11)) & (e11 = op1(e11, e11)) & (e11 = op1(e11, e11))) | ~ sP5), inference(nnf_transformation, [], [f28])).
fof(f2037, plain, (~ spl102_189 | ~ spl102_42), inference(avatar_split_clause, [], [f627, f1199, f2032])).
fof(f627, plain, (~ (e11 = op1(e11, e11)) | ~ sP5), inference(cnf_transformation, [], [f167])).
fof(f2030, plain, (~ spl102_188 | spl102_38), inference(avatar_split_clause, [], [f618, f1182, f2022])).
fof(f618, plain, ((e11 = op1(e11, e12)) | ~ sP6), inference(cnf_transformation, [], [f166])).
fof(f166, plain, ((~ (e13 = op1(e13, e11)) & ~ (e12 = op1(e12, e11)) & ~ (e11 = op1(e11, e11)) & ~ (e10 = op1(e10, e11)) & (e11 = op1(e12, e11)) & (e11 = op1(e11, e12))) | ~ sP6), inference(nnf_transformation, [], [f29])).
fof(f2029, plain, (~ spl102_188 | spl102_26), inference(avatar_split_clause, [], [f619, f1131, f2022])).
fof(f619, plain, ((e11 = op1(e12, e11)) | ~ sP6), inference(cnf_transformation, [], [f166])).
fof(f2027, plain, (~ spl102_188 | ~ spl102_42), inference(avatar_split_clause, [], [f621, f1199, f2022])).
fof(f621, plain, (~ (e11 = op1(e11, e11)) | ~ sP6), inference(cnf_transformation, [], [f166])).
fof(f2020, plain, (~ spl102_187 | spl102_34), inference(avatar_split_clause, [], [f612, f1165, f2012])).
fof(f612, plain, ((e11 = op1(e11, e13)) | ~ sP7), inference(cnf_transformation, [], [f165])).
fof(f165, plain, ((~ (e13 = op1(e13, e11)) & ~ (e12 = op1(e12, e11)) & ~ (e11 = op1(e11, e11)) & ~ (e10 = op1(e10, e11)) & (e11 = op1(e13, e11)) & (e11 = op1(e11, e13))) | ~ sP7), inference(nnf_transformation, [], [f30])).
fof(f2019, plain, (~ spl102_187 | spl102_10), inference(avatar_split_clause, [], [f613, f1063, f2012])).
fof(f613, plain, ((e11 = op1(e13, e11)) | ~ sP7), inference(cnf_transformation, [], [f165])).
fof(f2017, plain, (~ spl102_187 | ~ spl102_42), inference(avatar_split_clause, [], [f615, f1199, f2012])).
fof(f615, plain, (~ (e11 = op1(e11, e11)) | ~ sP7), inference(cnf_transformation, [], [f165])).
fof(f2016, plain, (~ spl102_187 | ~ spl102_27), inference(avatar_split_clause, [], [f616, f1135, f2012])).
fof(f616, plain, (~ (e12 = op1(e12, e11)) | ~ sP7), inference(cnf_transformation, [], [f165])).
fof(f2010, plain, (~ spl102_186 | spl102_31), inference(avatar_split_clause, [], [f606, f1152, f2002])).
fof(f606, plain, ((e12 = op1(e12, e10)) | ~ sP8), inference(cnf_transformation, [], [f164])).
fof(f164, plain, ((~ (e13 = op1(e13, e12)) & ~ (e12 = op1(e12, e12)) & ~ (e11 = op1(e11, e12)) & ~ (e10 = op1(e10, e12)) & (e12 = op1(e10, e12)) & (e12 = op1(e12, e10))) | ~ sP8), inference(nnf_transformation, [], [f31])).
fof(f2000, plain, (~ spl102_185 | spl102_27), inference(avatar_split_clause, [], [f600, f1135, f1992])).
fof(f600, plain, ((e12 = op1(e12, e11)) | ~ sP9), inference(cnf_transformation, [], [f163])).
fof(f163, plain, ((~ (e13 = op1(e13, e12)) & ~ (e12 = op1(e12, e12)) & ~ (e11 = op1(e11, e12)) & ~ (e10 = op1(e10, e12)) & (e12 = op1(e11, e12)) & (e12 = op1(e12, e11))) | ~ sP9), inference(nnf_transformation, [], [f32])).
fof(f1999, plain, (~ spl102_185 | spl102_39), inference(avatar_split_clause, [], [f601, f1186, f1992])).
fof(f601, plain, ((e12 = op1(e11, e12)) | ~ sP9), inference(cnf_transformation, [], [f163])).
fof(f1998, plain, (~ spl102_185 | ~ spl102_53), inference(avatar_split_clause, [], [f602, f1246, f1992])).
fof(f602, plain, (~ (e10 = op1(e10, e12)) | ~ sP9), inference(cnf_transformation, [], [f163])).
fof(f1997, plain, (~ spl102_185 | ~ spl102_38), inference(avatar_split_clause, [], [f603, f1182, f1992])).
fof(f603, plain, (~ (e11 = op1(e11, e12)) | ~ sP9), inference(cnf_transformation, [], [f163])).
fof(f1995, plain, (~ spl102_185 | ~ spl102_8), inference(avatar_split_clause, [], [f605, f1054, f1992])).
fof(f605, plain, (~ (e13 = op1(e13, e12)) | ~ sP9), inference(cnf_transformation, [], [f163])).
fof(f1990, plain, (~ spl102_184 | spl102_23), inference(avatar_split_clause, [], [f594, f1118, f1982])).
fof(f594, plain, ((e12 = op1(e12, e12)) | ~ sP10), inference(cnf_transformation, [], [f162])).
fof(f162, plain, ((~ (e13 = op1(e13, e12)) & ~ (e12 = op1(e12, e12)) & ~ (e11 = op1(e11, e12)) & ~ (e10 = op1(e10, e12)) & (e12 = op1(e12, e12)) & (e12 = op1(e12, e12))) | ~ sP10), inference(nnf_transformation, [], [f33])).
fof(f1986, plain, (~ spl102_184 | ~ spl102_23), inference(avatar_split_clause, [], [f598, f1118, f1982])).
fof(f598, plain, (~ (e12 = op1(e12, e12)) | ~ sP10), inference(cnf_transformation, [], [f162])).
fof(f1979, plain, (~ spl102_183 | spl102_7), inference(avatar_split_clause, [], [f589, f1050, f1972])).
fof(f589, plain, ((e12 = op1(e13, e12)) | ~ sP11), inference(cnf_transformation, [], [f161])).
fof(f161, plain, ((~ (e13 = op1(e13, e12)) & ~ (e12 = op1(e12, e12)) & ~ (e11 = op1(e11, e12)) & ~ (e10 = op1(e10, e12)) & (e12 = op1(e13, e12)) & (e12 = op1(e12, e13))) | ~ sP11), inference(nnf_transformation, [], [f34])).
fof(f1970, plain, (~ spl102_182 | spl102_16), inference(avatar_split_clause, [], [f582, f1088, f1962])).
fof(f582, plain, ((e13 = op1(e13, e10)) | ~ sP12), inference(cnf_transformation, [], [f160])).
fof(f160, plain, ((~ (e13 = op1(e13, e13)) & ~ (e12 = op1(e12, e13)) & ~ (e11 = op1(e11, e13)) & ~ (e10 = op1(e10, e13)) & (e13 = op1(e10, e13)) & (e13 = op1(e13, e10))) | ~ sP12), inference(nnf_transformation, [], [f35])).
fof(f1960, plain, (~ spl102_181 | spl102_12), inference(avatar_split_clause, [], [f576, f1071, f1952])).
fof(f576, plain, ((e13 = op1(e13, e11)) | ~ sP13), inference(cnf_transformation, [], [f159])).
fof(f159, plain, ((~ (e13 = op1(e13, e13)) & ~ (e12 = op1(e12, e13)) & ~ (e11 = op1(e11, e13)) & ~ (e10 = op1(e10, e13)) & (e13 = op1(e11, e13)) & (e13 = op1(e13, e11))) | ~ sP13), inference(nnf_transformation, [], [f36])).
fof(f1949, plain, (~ spl102_180 | spl102_20), inference(avatar_split_clause, [], [f571, f1105, f1942])).
fof(f571, plain, ((e13 = op1(e12, e13)) | ~ sP14), inference(cnf_transformation, [], [f158])).
fof(f158, plain, ((~ (e13 = op1(e13, e13)) & ~ (e12 = op1(e12, e13)) & ~ (e11 = op1(e11, e13)) & ~ (e10 = op1(e10, e13)) & (e13 = op1(e12, e13)) & (e13 = op1(e13, e12))) | ~ sP14), inference(nnf_transformation, [], [f37])).
fof(f1948, plain, (~ spl102_180 | ~ spl102_49), inference(avatar_split_clause, [], [f572, f1229, f1942])).
fof(f572, plain, (~ (e10 = op1(e10, e13)) | ~ sP14), inference(cnf_transformation, [], [f158])).
fof(f1947, plain, (~ spl102_180 | ~ spl102_34), inference(avatar_split_clause, [], [f573, f1165, f1942])).
fof(f573, plain, (~ (e11 = op1(e11, e13)) | ~ sP14), inference(cnf_transformation, [], [f158])).
fof(f1940, plain, ~ spl102_178, inference(avatar_split_clause, [], [f1009, f1932])).
fof(f1009, plain, ~ sP15, inference(trivial_inequality_removal, [], [f568])).
fof(f568, plain, (~ (op1(e10, e10) = op1(e10, e10)) | ~ sP15), inference(cnf_transformation, [], [f157])).
fof(f157, plain, (((e10 = op1(e10, op1(e10, e10))) & ~ (op1(e10, e10) = op1(e10, e10)) & (op1(e10, e10) = op1(e10, e10))) | ~ sP15), inference(nnf_transformation, [], [f38])).
fof(f1930, plain, (~ spl102_176 | spl102_171), inference(avatar_split_clause, [], [f564, f1897, f1922])).
fof(f564, plain, ((op1(e10, e11) = op1(e11, e10)) | ~ sP16), inference(cnf_transformation, [], [f156])).
fof(f156, plain, (((e11 = op1(e11, op1(e10, e11))) & ~ (op1(e10, e11) = op1(e11, e11)) & (op1(e10, e11) = op1(e11, e10))) | ~ sP16), inference(nnf_transformation, [], [f39])).
fof(f1920, plain, (~ spl102_174 | spl102_162), inference(avatar_split_clause, [], [f561, f1853, f1912])).
fof(f561, plain, ((op1(e10, e12) = op1(e12, e10)) | ~ sP17), inference(cnf_transformation, [], [f155])).
fof(f155, plain, (((e12 = op1(e12, op1(e10, e12))) & ~ (op1(e10, e12) = op1(e12, e12)) & (op1(e10, e12) = op1(e12, e10))) | ~ sP17), inference(nnf_transformation, [], [f40])).
fof(f1910, plain, (~ spl102_172 | spl102_152), inference(avatar_split_clause, [], [f558, f1805, f1902])).
fof(f558, plain, ((op1(e10, e13) = op1(e13, e10)) | ~ sP18), inference(cnf_transformation, [], [f154])).
fof(f154, plain, (((e13 = op1(e13, op1(e10, e13))) & ~ (op1(e10, e13) = op1(e13, e13)) & (op1(e10, e13) = op1(e13, e10))) | ~ sP18), inference(nnf_transformation, [], [f41])).
fof(f1900, plain, (~ spl102_169 | spl102_171), inference(avatar_split_clause, [], [f555, f1897, f1888])).
fof(f555, plain, ((op1(e10, e11) = op1(e11, e10)) | ~ sP19), inference(cnf_transformation, [], [f153])).
fof(f153, plain, (((e10 = op1(e10, op1(e11, e10))) & ~ (op1(e10, e10) = op1(e11, e10)) & (op1(e10, e11) = op1(e11, e10))) | ~ sP19), inference(nnf_transformation, [], [f42])).
fof(f1895, plain, (~ spl102_169 | spl102_170), inference(avatar_split_clause, [], [f557, f1892, f1888])).
fof(f557, plain, ((e10 = op1(e10, op1(e11, e10))) | ~ sP19), inference(cnf_transformation, [], [f153])).
fof(f1886, plain, ~ spl102_167, inference(avatar_split_clause, [], [f1010, f1878])).
fof(f1010, plain, ~ sP20, inference(trivial_inequality_removal, [], [f553])).
fof(f553, plain, (~ (op1(e11, e11) = op1(e11, e11)) | ~ sP20), inference(cnf_transformation, [], [f152])).
fof(f152, plain, (((e11 = op1(e11, op1(e11, e11))) & ~ (op1(e11, e11) = op1(e11, e11)) & (op1(e11, e11) = op1(e11, e11))) | ~ sP20), inference(nnf_transformation, [], [f43])).
fof(f1876, plain, (~ spl102_165 | spl102_159), inference(avatar_split_clause, [], [f549, f1839, f1868])).
fof(f549, plain, ((op1(e11, e12) = op1(e12, e11)) | ~ sP21), inference(cnf_transformation, [], [f151])).
fof(f151, plain, (((e12 = op1(e12, op1(e11, e12))) & ~ (op1(e11, e12) = op1(e12, e12)) & (op1(e11, e12) = op1(e12, e11))) | ~ sP21), inference(nnf_transformation, [], [f44])).
fof(f1875, plain, (~ spl102_165 | spl102_166), inference(avatar_split_clause, [], [f551, f1872, f1868])).
fof(f551, plain, ((e12 = op1(e12, op1(e11, e12))) | ~ sP21), inference(cnf_transformation, [], [f151])).
fof(f1866, plain, (~ spl102_163 | spl102_149), inference(avatar_split_clause, [], [f546, f1791, f1858])).
fof(f546, plain, ((op1(e11, e13) = op1(e13, e11)) | ~ sP22), inference(cnf_transformation, [], [f150])).
fof(f150, plain, (((e13 = op1(e13, op1(e11, e13))) & ~ (op1(e11, e13) = op1(e13, e13)) & (op1(e11, e13) = op1(e13, e11))) | ~ sP22), inference(nnf_transformation, [], [f45])).
fof(f1865, plain, (~ spl102_163 | spl102_164), inference(avatar_split_clause, [], [f548, f1862, f1858])).
fof(f548, plain, ((e13 = op1(e13, op1(e11, e13))) | ~ sP22), inference(cnf_transformation, [], [f150])).
fof(f1856, plain, (~ spl102_160 | spl102_162), inference(avatar_split_clause, [], [f543, f1853, f1844])).
fof(f543, plain, ((op1(e10, e12) = op1(e12, e10)) | ~ sP23), inference(cnf_transformation, [], [f149])).
fof(f149, plain, (((e10 = op1(e10, op1(e12, e10))) & ~ (op1(e10, e10) = op1(e12, e10)) & (op1(e10, e12) = op1(e12, e10))) | ~ sP23), inference(nnf_transformation, [], [f46])).
fof(f1842, plain, (~ spl102_157 | spl102_159), inference(avatar_split_clause, [], [f540, f1839, f1830])).
fof(f540, plain, ((op1(e11, e12) = op1(e12, e11)) | ~ sP24), inference(cnf_transformation, [], [f148])).
fof(f148, plain, (((e11 = op1(e11, op1(e12, e11))) & ~ (op1(e11, e11) = op1(e12, e11)) & (op1(e11, e12) = op1(e12, e11))) | ~ sP24), inference(nnf_transformation, [], [f47])).
fof(f1837, plain, (~ spl102_157 | spl102_158), inference(avatar_split_clause, [], [f542, f1834, f1830])).
fof(f542, plain, ((e11 = op1(e11, op1(e12, e11))) | ~ sP24), inference(cnf_transformation, [], [f148])).
fof(f1828, plain, ~ spl102_155, inference(avatar_split_clause, [], [f1011, f1820])).
fof(f1011, plain, ~ sP25, inference(trivial_inequality_removal, [], [f538])).
fof(f538, plain, (~ (op1(e12, e12) = op1(e12, e12)) | ~ sP25), inference(cnf_transformation, [], [f147])).
fof(f147, plain, (((e12 = op1(e12, op1(e12, e12))) & ~ (op1(e12, e12) = op1(e12, e12)) & (op1(e12, e12) = op1(e12, e12))) | ~ sP25), inference(nnf_transformation, [], [f48])).
fof(f1818, plain, (~ spl102_153 | spl102_146), inference(avatar_split_clause, [], [f534, f1777, f1810])).
fof(f534, plain, ((op1(e12, e13) = op1(e13, e12)) | ~ sP26), inference(cnf_transformation, [], [f146])).
fof(f146, plain, (((e13 = op1(e13, op1(e12, e13))) & ~ (op1(e12, e13) = op1(e13, e13)) & (op1(e12, e13) = op1(e13, e12))) | ~ sP26), inference(nnf_transformation, [], [f49])).
fof(f1808, plain, (~ spl102_150 | spl102_152), inference(avatar_split_clause, [], [f531, f1805, f1796])).
fof(f531, plain, ((op1(e10, e13) = op1(e13, e10)) | ~ sP27), inference(cnf_transformation, [], [f145])).
fof(f145, plain, (((e10 = op1(e10, op1(e13, e10))) & ~ (op1(e10, e10) = op1(e13, e10)) & (op1(e10, e13) = op1(e13, e10))) | ~ sP27), inference(nnf_transformation, [], [f50])).
fof(f1803, plain, (~ spl102_150 | spl102_151), inference(avatar_split_clause, [], [f533, f1800, f1796])).
fof(f533, plain, ((e10 = op1(e10, op1(e13, e10))) | ~ sP27), inference(cnf_transformation, [], [f145])).
fof(f1794, plain, (~ spl102_147 | spl102_149), inference(avatar_split_clause, [], [f528, f1791, f1782])).
fof(f528, plain, ((op1(e11, e13) = op1(e13, e11)) | ~ sP28), inference(cnf_transformation, [], [f144])).
fof(f144, plain, (((e11 = op1(e11, op1(e13, e11))) & ~ (op1(e11, e11) = op1(e13, e11)) & (op1(e11, e13) = op1(e13, e11))) | ~ sP28), inference(nnf_transformation, [], [f51])).
fof(f1789, plain, (~ spl102_147 | spl102_148), inference(avatar_split_clause, [], [f530, f1786, f1782])).
fof(f530, plain, ((e11 = op1(e11, op1(e13, e11))) | ~ sP28), inference(cnf_transformation, [], [f144])).
fof(f1780, plain, (~ spl102_144 | spl102_146), inference(avatar_split_clause, [], [f525, f1777, f1768])).
fof(f525, plain, ((op1(e12, e13) = op1(e13, e12)) | ~ sP29), inference(cnf_transformation, [], [f143])).
fof(f143, plain, (((e12 = op1(e12, op1(e13, e12))) & ~ (op1(e12, e12) = op1(e13, e12)) & (op1(e12, e13) = op1(e13, e12))) | ~ sP29), inference(nnf_transformation, [], [f52])).
fof(f1766, plain, (~ spl102_143 | spl102_61), inference(avatar_split_clause, [], [f520, f1280, f1759])).
fof(f520, plain, ((e10 = op1(e10, e10)) | ~ sP30), inference(cnf_transformation, [], [f142])).
fof(f142, plain, (((((~ (op1(e10, e10) = e13) | ~ (op1(e10, e10) = e13)) & (~ (op1(e10, e10) = e12) | ~ (op1(e10, e10) = e12)) & (~ (op1(e10, e10) = e11) | ~ (op1(e10, e10) = e11)) & (~ (e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10)))) | ~ (e10 = op1(e10, e10))) & (e10 = op1(e10, e10))) | ~ sP30), inference(nnf_transformation, [], [f53])).
fof(f1753, plain, (~ spl102_142 | ~ spl102_62 | ~ spl102_60 | ~ spl102_48), inference(avatar_split_clause, [], [f519, f1224, f1275, f1284, f1750])).
fof(f519, plain, (~ (e13 = op1(e11, e10)) | ~ (e13 = op1(e10, e11)) | ~ (op1(e10, e10) = e11) | ~ sP31), inference(cnf_transformation, [], [f141])).
fof(f141, plain, (((((~ (e13 = op1(e11, e10)) | ~ (e13 = op1(e10, e11))) & (~ (e12 = op1(e11, e10)) | ~ (e12 = op1(e10, e11))) & (~ (e11 = op1(e11, e10)) | ~ (e11 = op1(e10, e11))) & (~ (e10 = op1(e11, e10)) | ~ (e10 = op1(e10, e11)))) | ~ (op1(e10, e10) = e11)) & (op1(e10, e10) = e11)) | ~ sP31), inference(nnf_transformation, [], [f54])).
fof(f1748, plain, (~ spl102_141 | spl102_63), inference(avatar_split_clause, [], [f510, f1288, f1741])).
fof(f1288, plain, (spl102_63 <=> (op1(e10, e10) = e12)), introduced(avatar_definition, [new_symbols(naming, [spl102_63])])).
fof(f510, plain, ((op1(e10, e10) = e12) | ~ sP32), inference(cnf_transformation, [], [f140])).
fof(f140, plain, (((((~ (e13 = op1(e12, e10)) | ~ (e13 = op1(e10, e12))) & (~ (e12 = op1(e12, e10)) | ~ (e12 = op1(e10, e12))) & (~ (e11 = op1(e12, e10)) | ~ (e11 = op1(e10, e12))) & (~ (e10 = op1(e12, e10)) | ~ (e10 = op1(e10, e12)))) | ~ (op1(e10, e10) = e12)) & (op1(e10, e10) = e12)) | ~ sP32), inference(nnf_transformation, [], [f55])).
fof(f1747, plain, (~ spl102_141 | ~ spl102_63 | ~ spl102_53 | ~ spl102_29), inference(avatar_split_clause, [], [f511, f1144, f1246, f1288, f1741])).
fof(f511, plain, (~ (e10 = op1(e12, e10)) | ~ (e10 = op1(e10, e12)) | ~ (op1(e10, e10) = e12) | ~ sP32), inference(cnf_transformation, [], [f140])).
fof(f1739, plain, (~ spl102_140 | spl102_64), inference(avatar_split_clause, [], [f505, f1292, f1732])).
fof(f505, plain, ((op1(e10, e10) = e13) | ~ sP33), inference(cnf_transformation, [], [f139])).
fof(f139, plain, (((((~ (e13 = op1(e13, e10)) | ~ (e13 = op1(e10, e13))) & (~ (e12 = op1(e13, e10)) | ~ (e12 = op1(e10, e13))) & (~ (e11 = op1(e13, e10)) | ~ (e11 = op1(e10, e13))) & (~ (e10 = op1(e13, e10)) | ~ (e10 = op1(e10, e13)))) | ~ (op1(e10, e10) = e13)) & (op1(e10, e10) = e13)) | ~ sP33), inference(nnf_transformation, [], [f56])).
fof(f1730, plain, (~ spl102_139 | spl102_41), inference(avatar_split_clause, [], [f500, f1195, f1723])).
fof(f500, plain, ((e10 = op1(e11, e11)) | ~ sP34), inference(cnf_transformation, [], [f138])).
fof(f138, plain, (((((~ (e13 = op1(e10, e11)) | ~ (e13 = op1(e11, e10))) & (~ (e12 = op1(e10, e11)) | ~ (e12 = op1(e11, e10))) & (~ (e11 = op1(e10, e11)) | ~ (e11 = op1(e11, e10))) & (~ (e10 = op1(e10, e11)) | ~ (e10 = op1(e11, e10)))) | ~ (e10 = op1(e11, e11))) & (e10 = op1(e11, e11))) | ~ sP34), inference(nnf_transformation, [], [f57])).
fof(f1726, plain, (~ spl102_139 | ~ spl102_41 | ~ spl102_48 | ~ spl102_60), inference(avatar_split_clause, [], [f504, f1275, f1224, f1195, f1723])).
fof(f504, plain, (~ (e13 = op1(e10, e11)) | ~ (e13 = op1(e11, e10)) | ~ (e10 = op1(e11, e11)) | ~ sP34), inference(cnf_transformation, [], [f138])).
fof(f1721, plain, (~ spl102_138 | spl102_42), inference(avatar_split_clause, [], [f495, f1199, f1714])).
fof(f495, plain, ((e11 = op1(e11, e11)) | ~ sP35), inference(cnf_transformation, [], [f137])).
fof(f137, plain, (((((~ (e13 = op1(e11, e11)) | ~ (e13 = op1(e11, e11))) & (~ (e12 = op1(e11, e11)) | ~ (e12 = op1(e11, e11))) & (~ (e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & (~ (e10 = op1(e11, e11)) | ~ (e10 = op1(e11, e11)))) | ~ (e11 = op1(e11, e11))) & (e11 = op1(e11, e11))) | ~ sP35), inference(nnf_transformation, [], [f58])).
fof(f1719, plain, (~ spl102_138 | ~ spl102_42), inference(avatar_split_clause, [], [f1017, f1199, f1714])).
fof(f1017, plain, (~ (e11 = op1(e11, e11)) | ~ sP35), inference(duplicate_literal_removal, [], [f497])).
fof(f497, plain, (~ (e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11)) | ~ sP35), inference(cnf_transformation, [], [f137])).
fof(f1712, plain, (~ spl102_137 | spl102_43), inference(avatar_split_clause, [], [f490, f1203, f1705])).
fof(f490, plain, ((e12 = op1(e11, e11)) | ~ sP36), inference(cnf_transformation, [], [f136])).
fof(f136, plain, (((((~ (e13 = op1(e12, e11)) | ~ (e13 = op1(e11, e12))) & (~ (e12 = op1(e12, e11)) | ~ (e12 = op1(e11, e12))) & (~ (e11 = op1(e12, e11)) | ~ (e11 = op1(e11, e12))) & (~ (e10 = op1(e12, e11)) | ~ (e10 = op1(e11, e12)))) | ~ (e12 = op1(e11, e11))) & (e12 = op1(e11, e11))) | ~ sP36), inference(nnf_transformation, [], [f59])).
fof(f1710, plain, (~ spl102_137 | ~ spl102_43 | ~ spl102_38 | ~ spl102_26), inference(avatar_split_clause, [], [f492, f1131, f1182, f1203, f1705])).
fof(f492, plain, (~ (e11 = op1(e12, e11)) | ~ (e11 = op1(e11, e12)) | ~ (e12 = op1(e11, e11)) | ~ sP36), inference(cnf_transformation, [], [f136])).
fof(f1703, plain, (~ spl102_136 | spl102_44), inference(avatar_split_clause, [], [f485, f1207, f1696])).
fof(f485, plain, ((e13 = op1(e11, e11)) | ~ sP37), inference(cnf_transformation, [], [f135])).
fof(f135, plain, (((((~ (e13 = op1(e13, e11)) | ~ (e13 = op1(e11, e13))) & (~ (e12 = op1(e13, e11)) | ~ (e12 = op1(e11, e13))) & (~ (e11 = op1(e13, e11)) | ~ (e11 = op1(e11, e13))) & (~ (e10 = op1(e13, e11)) | ~ (e10 = op1(e11, e13)))) | ~ (e13 = op1(e11, e11))) & (e13 = op1(e11, e11))) | ~ sP37), inference(nnf_transformation, [], [f60])).
fof(f1694, plain, (~ spl102_135 | spl102_21), inference(avatar_split_clause, [], [f480, f1110, f1687])).
fof(f480, plain, ((e10 = op1(e12, e12)) | ~ sP38), inference(cnf_transformation, [], [f134])).
fof(f134, plain, (((((~ (e13 = op1(e10, e12)) | ~ (e13 = op1(e12, e10))) & (~ (e12 = op1(e10, e12)) | ~ (e12 = op1(e12, e10))) & (~ (e11 = op1(e10, e12)) | ~ (e11 = op1(e12, e10))) & (~ (e10 = op1(e10, e12)) | ~ (e10 = op1(e12, e10)))) | ~ (e10 = op1(e12, e12))) & (e10 = op1(e12, e12))) | ~ sP38), inference(nnf_transformation, [], [f61])).
fof(f1685, plain, (~ spl102_134 | spl102_22), inference(avatar_split_clause, [], [f475, f1114, f1678])).
fof(f475, plain, ((e11 = op1(e12, e12)) | ~ sP39), inference(cnf_transformation, [], [f133])).
fof(f133, plain, (((((~ (e13 = op1(e11, e12)) | ~ (e13 = op1(e12, e11))) & (~ (e12 = op1(e11, e12)) | ~ (e12 = op1(e12, e11))) & (~ (e11 = op1(e11, e12)) | ~ (e11 = op1(e12, e11))) & (~ (e10 = op1(e11, e12)) | ~ (e10 = op1(e12, e11)))) | ~ (e11 = op1(e12, e12))) & (e11 = op1(e12, e12))) | ~ sP39), inference(nnf_transformation, [], [f62])).
fof(f1682, plain, (~ spl102_134 | ~ spl102_22 | ~ spl102_27 | ~ spl102_39), inference(avatar_split_clause, [], [f478, f1186, f1135, f1114, f1678])).
fof(f478, plain, (~ (e12 = op1(e11, e12)) | ~ (e12 = op1(e12, e11)) | ~ (e11 = op1(e12, e12)) | ~ sP39), inference(cnf_transformation, [], [f133])).
fof(f1676, plain, (~ spl102_133 | spl102_23), inference(avatar_split_clause, [], [f470, f1118, f1669])).
fof(f470, plain, ((e12 = op1(e12, e12)) | ~ sP40), inference(cnf_transformation, [], [f132])).
fof(f132, plain, (((((~ (e13 = op1(e12, e12)) | ~ (e13 = op1(e12, e12))) & (~ (e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & (~ (e11 = op1(e12, e12)) | ~ (e11 = op1(e12, e12))) & (~ (e10 = op1(e12, e12)) | ~ (e10 = op1(e12, e12)))) | ~ (e12 = op1(e12, e12))) & (e12 = op1(e12, e12))) | ~ sP40), inference(nnf_transformation, [], [f63])).
fof(f1673, plain, (~ spl102_133 | ~ spl102_23), inference(avatar_split_clause, [], [f1022, f1118, f1669])).
fof(f1022, plain, (~ (e12 = op1(e12, e12)) | ~ sP40), inference(duplicate_literal_removal, [], [f473])).
fof(f473, plain, (~ (e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12)) | ~ sP40), inference(cnf_transformation, [], [f132])).
fof(f1667, plain, (~ spl102_132 | spl102_24), inference(avatar_split_clause, [], [f465, f1122, f1660])).
fof(f465, plain, ((e13 = op1(e12, e12)) | ~ sP41), inference(cnf_transformation, [], [f131])).
fof(f131, plain, (((((~ (e13 = op1(e13, e12)) | ~ (e13 = op1(e12, e13))) & (~ (e12 = op1(e13, e12)) | ~ (e12 = op1(e12, e13))) & (~ (e11 = op1(e13, e12)) | ~ (e11 = op1(e12, e13))) & (~ (e10 = op1(e13, e12)) | ~ (e10 = op1(e12, e13)))) | ~ (e13 = op1(e12, e12))) & (e13 = op1(e12, e12))) | ~ sP41), inference(nnf_transformation, [], [f64])).
fof(f1665, plain, (~ spl102_132 | ~ spl102_24 | ~ spl102_18 | ~ spl102_6), inference(avatar_split_clause, [], [f467, f1046, f1097, f1122, f1660])).
fof(f467, plain, (~ (e11 = op1(e13, e12)) | ~ (e11 = op1(e12, e13)) | ~ (e13 = op1(e12, e12)) | ~ sP41), inference(cnf_transformation, [], [f131])).
fof(f1658, plain, (~ spl102_131 | spl102_1), inference(avatar_split_clause, [], [f460, f1025, f1651])).
fof(f460, plain, ((e10 = op1(e13, e13)) | ~ sP42), inference(cnf_transformation, [], [f130])).
fof(f130, plain, (((((~ (e13 = op1(e10, e13)) | ~ (e13 = op1(e13, e10))) & (~ (e12 = op1(e10, e13)) | ~ (e12 = op1(e13, e10))) & (~ (e11 = op1(e10, e13)) | ~ (e11 = op1(e13, e10))) & (~ (e10 = op1(e10, e13)) | ~ (e10 = op1(e13, e10)))) | ~ (e10 = op1(e13, e13))) & (e10 = op1(e13, e13))) | ~ sP42), inference(nnf_transformation, [], [f65])).
fof(f1655, plain, (~ spl102_131 | ~ spl102_1 | ~ spl102_15 | ~ spl102_51), inference(avatar_split_clause, [], [f463, f1237, f1084, f1025, f1651])).
fof(f463, plain, (~ (e12 = op1(e10, e13)) | ~ (e12 = op1(e13, e10)) | ~ (e10 = op1(e13, e13)) | ~ sP42), inference(cnf_transformation, [], [f130])).
fof(f1649, plain, (~ spl102_130 | spl102_2), inference(avatar_split_clause, [], [f455, f1029, f1642])).
fof(f455, plain, ((e11 = op1(e13, e13)) | ~ sP43), inference(cnf_transformation, [], [f129])).
fof(f129, plain, (((((~ (e13 = op1(e11, e13)) | ~ (e13 = op1(e13, e11))) & (~ (e12 = op1(e11, e13)) | ~ (e12 = op1(e13, e11))) & (~ (e11 = op1(e11, e13)) | ~ (e11 = op1(e13, e11))) & (~ (e10 = op1(e11, e13)) | ~ (e10 = op1(e13, e11)))) | ~ (e11 = op1(e13, e13))) & (e11 = op1(e13, e13))) | ~ sP43), inference(nnf_transformation, [], [f66])).
fof(f1648, plain, (~ spl102_130 | ~ spl102_2 | ~ spl102_9 | ~ spl102_33), inference(avatar_split_clause, [], [f456, f1161, f1059, f1029, f1642])).
fof(f456, plain, (~ (e10 = op1(e11, e13)) | ~ (e10 = op1(e13, e11)) | ~ (e11 = op1(e13, e13)) | ~ sP43), inference(cnf_transformation, [], [f129])).
fof(f1640, plain, (~ spl102_129 | spl102_3), inference(avatar_split_clause, [], [f450, f1033, f1633])).
fof(f450, plain, ((e12 = op1(e13, e13)) | ~ sP44), inference(cnf_transformation, [], [f128])).
fof(f128, plain, (((((~ (e13 = op1(e12, e13)) | ~ (e13 = op1(e13, e12))) & (~ (e12 = op1(e12, e13)) | ~ (e12 = op1(e13, e12))) & (~ (e11 = op1(e12, e13)) | ~ (e11 = op1(e13, e12))) & (~ (e10 = op1(e12, e13)) | ~ (e10 = op1(e13, e12)))) | ~ (e12 = op1(e13, e13))) & (e12 = op1(e13, e13))) | ~ sP44), inference(nnf_transformation, [], [f67])).
fof(f1631, plain, (spl102_125 | spl102_121 | spl102_117 | spl102_113), inference(avatar_split_clause, [], [f294, f1533, f1550, f1567, f1584])).
fof(f294, plain, ((e20 = op2(e20, e23)) | (e20 = op2(e20, e22)) | (e20 = op2(e20, e21)) | (e20 = op2(e20, e20))), inference(cnf_transformation, [], [f4])).
fof(f4, plain, (((e23 = op2(e23, e23)) | (e23 = op2(e22, e23)) | (e23 = op2(e21, e23)) | (e23 = op2(e20, e23))) & ((e23 = op2(e23, e23)) | (e23 = op2(e23, e22)) | (e23 = op2(e23, e21)) | (e23 = op2(e23, e20))) & ((e22 = op2(e23, e23)) | (e22 = op2(e22, e23)) | (e22 = op2(e21, e23)) | (e22 = op2(e20, e23))) & ((e22 = op2(e23, e23)) | (e22 = op2(e23, e22)) | (e22 = op2(e23, e21)) | (e22 = op2(e23, e20))) & ((e21 = op2(e23, e23)) | (e21 = op2(e22, e23)) | (e21 = op2(e21, e23)) | (e21 = op2(e20, e23))) & ((e21 = op2(e23, e23)) | (e21 = op2(e23, e22)) | (e21 = op2(e23, e21)) | (e21 = op2(e23, e20))) & ((e20 = op2(e23, e23)) | (e20 = op2(e22, e23)) | (e20 = op2(e21, e23)) | (e20 = op2(e20, e23))) & ((e20 = op2(e23, e23)) | (e20 = op2(e23, e22)) | (e20 = op2(e23, e21)) | (e20 = op2(e23, e20))) & ((e23 = op2(e23, e22)) | (e23 = op2(e22, e22)) | (e23 = op2(e21, e22)) | (e23 = op2(e20, e22))) & ((e23 = op2(e22, e23)) | (e23 = op2(e22, e22)) | (e23 = op2(e22, e21)) | (e23 = op2(e22, e20))) & ((e22 = op2(e23, e22)) | (e22 = op2(e22, e22)) | (e22 = op2(e21, e22)) | (e22 = op2(e20, e22))) & ((e22 = op2(e22, e23)) | (e22 = op2(e22, e22)) | (e22 = op2(e22, e21)) | (e22 = op2(e22, e20))) & ((e21 = op2(e23, e22)) | (e21 = op2(e22, e22)) | (e21 = op2(e21, e22)) | (e21 = op2(e20, e22))) & ((e21 = op2(e22, e23)) | (e21 = op2(e22, e22)) | (e21 = op2(e22, e21)) | (e21 = op2(e22, e20))) & ((e20 = op2(e23, e22)) | (e20 = op2(e22, e22)) | (e20 = op2(e21, e22)) | (e20 = op2(e20, e22))) & ((e20 = op2(e22, e23)) | (e20 = op2(e22, e22)) | (e20 = op2(e22, e21)) | (e20 = op2(e22, e20))) & ((e23 = op2(e23, e21)) | (e23 = op2(e22, e21)) | (e23 = op2(e21, e21)) | (e23 = op2(e20, e21))) & ((e23 = op2(e21, e23)) | (e23 = op2(e21, e22)) | (e23 = op2(e21, e21)) | (e23 = op2(e21, e20))) & ((e22 = op2(e23, e21)) | (e22 = op2(e22, e21)) | (e22 = op2(e21, e21)) | (e22 = op2(e20, e21))) & ((e22 = op2(e21, e23)) | (e22 = op2(e21, e22)) | (e22 = op2(e21, e21)) | (e22 = op2(e21, e20))) & ((e21 = op2(e23, e21)) | (e21 = op2(e22, e21)) | (e21 = op2(e21, e21)) | (e21 = op2(e20, e21))) & ((e21 = op2(e21, e23)) | (e21 = op2(e21, e22)) | (e21 = op2(e21, e21)) | (e21 = op2(e21, e20))) & ((e20 = op2(e23, e21)) | (e20 = op2(e22, e21)) | (e20 = op2(e21, e21)) | (e20 = op2(e20, e21))) & ((e20 = op2(e21, e23)) | (e20 = op2(e21, e22)) | (e20 = op2(e21, e21)) | (e20 = op2(e21, e20))) & ((e23 = op2(e23, e20)) | (e23 = op2(e22, e20)) | (e23 = op2(e21, e20)) | (op2(e20, e20) = e23)) & ((e23 = op2(e20, e23)) | (e23 = op2(e20, e22)) | (e23 = op2(e20, e21)) | (op2(e20, e20) = e23)) & ((e22 = op2(e23, e20)) | (e22 = op2(e22, e20)) | (e22 = op2(e21, e20)) | (op2(e20, e20) = e22)) & ((e22 = op2(e20, e23)) | (e22 = op2(e20, e22)) | (e22 = op2(e20, e21)) | (op2(e20, e20) = e22)) & ((e21 = op2(e23, e20)) | (e21 = op2(e22, e20)) | (e21 = op2(e21, e20)) | (op2(e20, e20) = e21)) & ((e21 = op2(e20, e23)) | (e21 = op2(e20, e22)) | (e21 = op2(e20, e21)) | (op2(e20, e20) = e21)) & ((e20 = op2(e23, e20)) | (e20 = op2(e22, e20)) | (e20 = op2(e21, e20)) | (e20 = op2(e20, e20))) & ((e20 = op2(e20, e23)) | (e20 = op2(e20, e22)) | (e20 = op2(e20, e21)) | (e20 = op2(e20, e20)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG109+1.p', ax4)).
fof(f1630, plain, (spl102_125 | spl102_109 | spl102_93 | spl102_77), inference(avatar_split_clause, [], [f295, f1380, f1448, f1516, f1584])).
fof(f295, plain, ((e20 = op2(e23, e20)) | (e20 = op2(e22, e20)) | (e20 = op2(e21, e20)) | (e20 = op2(e20, e20))), inference(cnf_transformation, [], [f4])).
fof(f1624, plain, (spl102_128 | spl102_112 | spl102_96 | spl102_80), inference(avatar_split_clause, [], [f301, f1392, f1460, f1528, f1596])).
fof(f301, plain, ((e23 = op2(e23, e20)) | (e23 = op2(e22, e20)) | (e23 = op2(e21, e20)) | (op2(e20, e20) = e23)), inference(cnf_transformation, [], [f4])).
fof(f1623, plain, (spl102_109 | spl102_105 | spl102_101 | spl102_97), inference(avatar_split_clause, [], [f302, f1465, f1482, f1499, f1516])).
fof(f302, plain, ((e20 = op2(e21, e23)) | (e20 = op2(e21, e22)) | (e20 = op2(e21, e21)) | (e20 = op2(e21, e20))), inference(cnf_transformation, [], [f4])).
fof(f1622, plain, (spl102_121 | spl102_105 | spl102_89 | spl102_73), inference(avatar_split_clause, [], [f303, f1363, f1431, f1499, f1567])).
fof(f303, plain, ((e20 = op2(e23, e21)) | (e20 = op2(e22, e21)) | (e20 = op2(e21, e21)) | (e20 = op2(e20, e21))), inference(cnf_transformation, [], [f4])).
fof(f1620, plain, (spl102_122 | spl102_106 | spl102_90 | spl102_74), inference(avatar_split_clause, [], [f305, f1367, f1435, f1503, f1571])).
fof(f305, plain, ((e21 = op2(e23, e21)) | (e21 = op2(e22, e21)) | (e21 = op2(e21, e21)) | (e21 = op2(e20, e21))), inference(cnf_transformation, [], [f4])).
fof(f1619, plain, (spl102_111 | spl102_107 | spl102_103 | spl102_99), inference(avatar_split_clause, [], [f306, f1473, f1490, f1507, f1524])).
fof(f306, plain, ((e22 = op2(e21, e23)) | (e22 = op2(e21, e22)) | (e22 = op2(e21, e21)) | (e22 = op2(e21, e20))), inference(cnf_transformation, [], [f4])).
fof(f1617, plain, (spl102_112 | spl102_108 | spl102_104 | spl102_100), inference(avatar_split_clause, [], [f308, f1477, f1494, f1511, f1528])).
fof(f308, plain, ((e23 = op2(e21, e23)) | (e23 = op2(e21, e22)) | (e23 = op2(e21, e21)) | (e23 = op2(e21, e20))), inference(cnf_transformation, [], [f4])).
fof(f1612, plain, (spl102_118 | spl102_102 | spl102_86 | spl102_70), inference(avatar_split_clause, [], [f313, f1350, f1418, f1486, f1554])).
fof(f313, plain, ((e21 = op2(e23, e22)) | (e21 = op2(e22, e22)) | (e21 = op2(e21, e22)) | (e21 = op2(e20, e22))), inference(cnf_transformation, [], [f4])).
fof(f1611, plain, (spl102_95 | spl102_91 | spl102_87 | spl102_83), inference(avatar_split_clause, [], [f314, f1405, f1422, f1439, f1456])).
fof(f314, plain, ((e22 = op2(e22, e23)) | (e22 = op2(e22, e22)) | (e22 = op2(e22, e21)) | (e22 = op2(e22, e20))), inference(cnf_transformation, [], [f4])).
fof(f1610, plain, (spl102_119 | spl102_103 | spl102_87 | spl102_71), inference(avatar_split_clause, [], [f315, f1354, f1422, f1490, f1558])).
fof(f315, plain, ((e22 = op2(e23, e22)) | (e22 = op2(e22, e22)) | (e22 = op2(e21, e22)) | (e22 = op2(e20, e22))), inference(cnf_transformation, [], [f4])).
fof(f1608, plain, (spl102_120 | spl102_104 | spl102_88 | spl102_72), inference(avatar_split_clause, [], [f317, f1358, f1426, f1494, f1562])).
fof(f317, plain, ((e23 = op2(e23, e22)) | (e23 = op2(e22, e22)) | (e23 = op2(e21, e22)) | (e23 = op2(e20, e22))), inference(cnf_transformation, [], [f4])).
fof(f1607, plain, (spl102_77 | spl102_73 | spl102_69 | spl102_65), inference(avatar_split_clause, [], [f318, f1329, f1346, f1363, f1380])).
fof(f318, plain, ((e20 = op2(e23, e23)) | (e20 = op2(e23, e22)) | (e20 = op2(e23, e21)) | (e20 = op2(e23, e20))), inference(cnf_transformation, [], [f4])).
fof(f1606, plain, (spl102_113 | spl102_97 | spl102_81 | spl102_65), inference(avatar_split_clause, [], [f319, f1329, f1397, f1465, f1533])).
fof(f319, plain, ((e20 = op2(e23, e23)) | (e20 = op2(e22, e23)) | (e20 = op2(e21, e23)) | (e20 = op2(e20, e23))), inference(cnf_transformation, [], [f4])).
fof(f1604, plain, (spl102_114 | spl102_98 | spl102_82 | spl102_66), inference(avatar_split_clause, [], [f321, f1333, f1401, f1469, f1537])).
fof(f321, plain, ((e21 = op2(e23, e23)) | (e21 = op2(e22, e23)) | (e21 = op2(e21, e23)) | (e21 = op2(e20, e23))), inference(cnf_transformation, [], [f4])).
fof(f1602, plain, (spl102_115 | spl102_99 | spl102_83 | spl102_67), inference(avatar_split_clause, [], [f323, f1337, f1405, f1473, f1541])).
fof(f323, plain, ((e22 = op2(e23, e23)) | (e22 = op2(e22, e23)) | (e22 = op2(e21, e23)) | (e22 = op2(e20, e23))), inference(cnf_transformation, [], [f4])).
fof(f1601, plain, (spl102_80 | spl102_76 | spl102_72 | spl102_68), inference(avatar_split_clause, [], [f324, f1341, f1358, f1375, f1392])).
fof(f324, plain, ((e23 = op2(e23, e23)) | (e23 = op2(e23, e22)) | (e23 = op2(e23, e21)) | (e23 = op2(e23, e20))), inference(cnf_transformation, [], [f4])).
fof(f1565, plain, (spl102_117 | spl102_118 | spl102_119 | spl102_120), inference(avatar_split_clause, [], [f280, f1562, f1558, f1554, f1550])).
fof(f280, plain, ((e23 = op2(e20, e22)) | (e22 = op2(e20, e22)) | (e21 = op2(e20, e22)) | (e20 = op2(e20, e22))), inference(cnf_transformation, [], [f3])).
fof(f3, plain, (((e23 = op2(e23, e23)) | (e22 = op2(e23, e23)) | (e21 = op2(e23, e23)) | (e20 = op2(e23, e23))) & ((e23 = op2(e23, e22)) | (e22 = op2(e23, e22)) | (e21 = op2(e23, e22)) | (e20 = op2(e23, e22))) & ((e23 = op2(e23, e21)) | (e22 = op2(e23, e21)) | (e21 = op2(e23, e21)) | (e20 = op2(e23, e21))) & ((e23 = op2(e23, e20)) | (e22 = op2(e23, e20)) | (e21 = op2(e23, e20)) | (e20 = op2(e23, e20))) & ((e23 = op2(e22, e23)) | (e22 = op2(e22, e23)) | (e21 = op2(e22, e23)) | (e20 = op2(e22, e23))) & ((e23 = op2(e22, e22)) | (e22 = op2(e22, e22)) | (e21 = op2(e22, e22)) | (e20 = op2(e22, e22))) & ((e23 = op2(e22, e21)) | (e22 = op2(e22, e21)) | (e21 = op2(e22, e21)) | (e20 = op2(e22, e21))) & ((e23 = op2(e22, e20)) | (e22 = op2(e22, e20)) | (e21 = op2(e22, e20)) | (e20 = op2(e22, e20))) & ((e23 = op2(e21, e23)) | (e22 = op2(e21, e23)) | (e21 = op2(e21, e23)) | (e20 = op2(e21, e23))) & ((e23 = op2(e21, e22)) | (e22 = op2(e21, e22)) | (e21 = op2(e21, e22)) | (e20 = op2(e21, e22))) & ((e23 = op2(e21, e21)) | (e22 = op2(e21, e21)) | (e21 = op2(e21, e21)) | (e20 = op2(e21, e21))) & ((e23 = op2(e21, e20)) | (e22 = op2(e21, e20)) | (e21 = op2(e21, e20)) | (e20 = op2(e21, e20))) & ((e23 = op2(e20, e23)) | (e22 = op2(e20, e23)) | (e21 = op2(e20, e23)) | (e20 = op2(e20, e23))) & ((e23 = op2(e20, e22)) | (e22 = op2(e20, e22)) | (e21 = op2(e20, e22)) | (e20 = op2(e20, e22))) & ((e23 = op2(e20, e21)) | (e22 = op2(e20, e21)) | (e21 = op2(e20, e21)) | (e20 = op2(e20, e21))) & ((op2(e20, e20) = e23) | (op2(e20, e20) = e22) | (op2(e20, e20) = e21) | (e20 = op2(e20, e20)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG109+1.p', ax3)).
fof(f1548, plain, (spl102_113 | spl102_114 | spl102_115 | spl102_116), inference(avatar_split_clause, [], [f281, f1545, f1541, f1537, f1533])).
fof(f281, plain, ((e23 = op2(e20, e23)) | (e22 = op2(e20, e23)) | (e21 = op2(e20, e23)) | (e20 = op2(e20, e23))), inference(cnf_transformation, [], [f3])).
fof(f1531, plain, (spl102_109 | spl102_110 | spl102_111 | spl102_112), inference(avatar_split_clause, [], [f282, f1528, f1524, f1520, f1516])).
fof(f282, plain, ((e23 = op2(e21, e20)) | (e22 = op2(e21, e20)) | (e21 = op2(e21, e20)) | (e20 = op2(e21, e20))), inference(cnf_transformation, [], [f3])).
fof(f1514, plain, (spl102_105 | spl102_106 | spl102_107 | spl102_108), inference(avatar_split_clause, [], [f283, f1511, f1507, f1503, f1499])).
fof(f283, plain, ((e23 = op2(e21, e21)) | (e22 = op2(e21, e21)) | (e21 = op2(e21, e21)) | (e20 = op2(e21, e21))), inference(cnf_transformation, [], [f3])).
fof(f1497, plain, (spl102_101 | spl102_102 | spl102_103 | spl102_104), inference(avatar_split_clause, [], [f284, f1494, f1490, f1486, f1482])).
fof(f284, plain, ((e23 = op2(e21, e22)) | (e22 = op2(e21, e22)) | (e21 = op2(e21, e22)) | (e20 = op2(e21, e22))), inference(cnf_transformation, [], [f3])).
fof(f1480, plain, (spl102_97 | spl102_98 | spl102_99 | spl102_100), inference(avatar_split_clause, [], [f285, f1477, f1473, f1469, f1465])).
fof(f285, plain, ((e23 = op2(e21, e23)) | (e22 = op2(e21, e23)) | (e21 = op2(e21, e23)) | (e20 = op2(e21, e23))), inference(cnf_transformation, [], [f3])).
fof(f1463, plain, (spl102_93 | spl102_94 | spl102_95 | spl102_96), inference(avatar_split_clause, [], [f286, f1460, f1456, f1452, f1448])).
fof(f286, plain, ((e23 = op2(e22, e20)) | (e22 = op2(e22, e20)) | (e21 = op2(e22, e20)) | (e20 = op2(e22, e20))), inference(cnf_transformation, [], [f3])).
fof(f1446, plain, (spl102_89 | spl102_90 | spl102_91 | spl102_92), inference(avatar_split_clause, [], [f287, f1443, f1439, f1435, f1431])).
fof(f287, plain, ((e23 = op2(e22, e21)) | (e22 = op2(e22, e21)) | (e21 = op2(e22, e21)) | (e20 = op2(e22, e21))), inference(cnf_transformation, [], [f3])).
fof(f1429, plain, (spl102_85 | spl102_86 | spl102_87 | spl102_88), inference(avatar_split_clause, [], [f288, f1426, f1422, f1418, f1414])).
fof(f288, plain, ((e23 = op2(e22, e22)) | (e22 = op2(e22, e22)) | (e21 = op2(e22, e22)) | (e20 = op2(e22, e22))), inference(cnf_transformation, [], [f3])).
fof(f1378, plain, (spl102_73 | spl102_74 | spl102_75 | spl102_76), inference(avatar_split_clause, [], [f291, f1375, f1371, f1367, f1363])).
fof(f291, plain, ((e23 = op2(e23, e21)) | (e22 = op2(e23, e21)) | (e21 = op2(e23, e21)) | (e20 = op2(e23, e21))), inference(cnf_transformation, [], [f3])).
fof(f1361, plain, (spl102_69 | spl102_70 | spl102_71 | spl102_72), inference(avatar_split_clause, [], [f292, f1358, f1354, f1350, f1346])).
fof(f292, plain, ((e23 = op2(e23, e22)) | (e22 = op2(e23, e22)) | (e21 = op2(e23, e22)) | (e20 = op2(e23, e22))), inference(cnf_transformation, [], [f3])).
fof(f1327, plain, (spl102_61 | spl102_57 | spl102_53 | spl102_49), inference(avatar_split_clause, [], [f246, f1229, f1246, f1263, f1280])).
fof(f246, plain, ((e10 = op1(e10, e13)) | (e10 = op1(e10, e12)) | (e10 = op1(e10, e11)) | (e10 = op1(e10, e10))), inference(cnf_transformation, [], [f2])).
fof(f2, plain, (((e13 = op1(e13, e13)) | (e13 = op1(e12, e13)) | (e13 = op1(e11, e13)) | (e13 = op1(e10, e13))) & ((e13 = op1(e13, e13)) | (e13 = op1(e13, e12)) | (e13 = op1(e13, e11)) | (e13 = op1(e13, e10))) & ((e12 = op1(e13, e13)) | (e12 = op1(e12, e13)) | (e12 = op1(e11, e13)) | (e12 = op1(e10, e13))) & ((e12 = op1(e13, e13)) | (e12 = op1(e13, e12)) | (e12 = op1(e13, e11)) | (e12 = op1(e13, e10))) & ((e11 = op1(e13, e13)) | (e11 = op1(e12, e13)) | (e11 = op1(e11, e13)) | (e11 = op1(e10, e13))) & ((e11 = op1(e13, e13)) | (e11 = op1(e13, e12)) | (e11 = op1(e13, e11)) | (e11 = op1(e13, e10))) & ((e10 = op1(e13, e13)) | (e10 = op1(e12, e13)) | (e10 = op1(e11, e13)) | (e10 = op1(e10, e13))) & ((e10 = op1(e13, e13)) | (e10 = op1(e13, e12)) | (e10 = op1(e13, e11)) | (e10 = op1(e13, e10))) & ((e13 = op1(e13, e12)) | (e13 = op1(e12, e12)) | (e13 = op1(e11, e12)) | (e13 = op1(e10, e12))) & ((e13 = op1(e12, e13)) | (e13 = op1(e12, e12)) | (e13 = op1(e12, e11)) | (e13 = op1(e12, e10))) & ((e12 = op1(e13, e12)) | (e12 = op1(e12, e12)) | (e12 = op1(e11, e12)) | (e12 = op1(e10, e12))) & ((e12 = op1(e12, e13)) | (e12 = op1(e12, e12)) | (e12 = op1(e12, e11)) | (e12 = op1(e12, e10))) & ((e11 = op1(e13, e12)) | (e11 = op1(e12, e12)) | (e11 = op1(e11, e12)) | (e11 = op1(e10, e12))) & ((e11 = op1(e12, e13)) | (e11 = op1(e12, e12)) | (e11 = op1(e12, e11)) | (e11 = op1(e12, e10))) & ((e10 = op1(e13, e12)) | (e10 = op1(e12, e12)) | (e10 = op1(e11, e12)) | (e10 = op1(e10, e12))) & ((e10 = op1(e12, e13)) | (e10 = op1(e12, e12)) | (e10 = op1(e12, e11)) | (e10 = op1(e12, e10))) & ((e13 = op1(e13, e11)) | (e13 = op1(e12, e11)) | (e13 = op1(e11, e11)) | (e13 = op1(e10, e11))) & ((e13 = op1(e11, e13)) | (e13 = op1(e11, e12)) | (e13 = op1(e11, e11)) | (e13 = op1(e11, e10))) & ((e12 = op1(e13, e11)) | (e12 = op1(e12, e11)) | (e12 = op1(e11, e11)) | (e12 = op1(e10, e11))) & ((e12 = op1(e11, e13)) | (e12 = op1(e11, e12)) | (e12 = op1(e11, e11)) | (e12 = op1(e11, e10))) & ((e11 = op1(e13, e11)) | (e11 = op1(e12, e11)) | (e11 = op1(e11, e11)) | (e11 = op1(e10, e11))) & ((e11 = op1(e11, e13)) | (e11 = op1(e11, e12)) | (e11 = op1(e11, e11)) | (e11 = op1(e11, e10))) & ((e10 = op1(e13, e11)) | (e10 = op1(e12, e11)) | (e10 = op1(e11, e11)) | (e10 = op1(e10, e11))) & ((e10 = op1(e11, e13)) | (e10 = op1(e11, e12)) | (e10 = op1(e11, e11)) | (e10 = op1(e11, e10))) & ((e13 = op1(e13, e10)) | (e13 = op1(e12, e10)) | (e13 = op1(e11, e10)) | (op1(e10, e10) = e13)) & ((e13 = op1(e10, e13)) | (e13 = op1(e10, e12)) | (e13 = op1(e10, e11)) | (op1(e10, e10) = e13)) & ((e12 = op1(e13, e10)) | (e12 = op1(e12, e10)) | (e12 = op1(e11, e10)) | (op1(e10, e10) = e12)) & ((e12 = op1(e10, e13)) | (e12 = op1(e10, e12)) | (e12 = op1(e10, e11)) | (op1(e10, e10) = e12)) & ((e11 = op1(e13, e10)) | (e11 = op1(e12, e10)) | (e11 = op1(e11, e10)) | (op1(e10, e10) = e11)) & ((e11 = op1(e10, e13)) | (e11 = op1(e10, e12)) | (e11 = op1(e10, e11)) | (op1(e10, e10) = e11)) & ((e10 = op1(e13, e10)) | (e10 = op1(e12, e10)) | (e10 = op1(e11, e10)) | (e10 = op1(e10, e10))) & ((e10 = op1(e10, e13)) | (e10 = op1(e10, e12)) | (e10 = op1(e10, e11)) | (e10 = op1(e10, e10)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG109+1.p', ax2)).
fof(f1326, plain, (spl102_61 | spl102_45 | spl102_29 | spl102_13), inference(avatar_split_clause, [], [f247, f1076, f1144, f1212, f1280])).
fof(f247, plain, ((e10 = op1(e13, e10)) | (e10 = op1(e12, e10)) | (e10 = op1(e11, e10)) | (e10 = op1(e10, e10))), inference(cnf_transformation, [], [f2])).
fof(f1320, plain, (spl102_64 | spl102_48 | spl102_32 | spl102_16), inference(avatar_split_clause, [], [f253, f1088, f1156, f1224, f1292])).
fof(f253, plain, ((e13 = op1(e13, e10)) | (e13 = op1(e12, e10)) | (e13 = op1(e11, e10)) | (op1(e10, e10) = e13)), inference(cnf_transformation, [], [f2])).
fof(f1319, plain, (spl102_45 | spl102_41 | spl102_37 | spl102_33), inference(avatar_split_clause, [], [f254, f1161, f1178, f1195, f1212])).
fof(f254, plain, ((e10 = op1(e11, e13)) | (e10 = op1(e11, e12)) | (e10 = op1(e11, e11)) | (e10 = op1(e11, e10))), inference(cnf_transformation, [], [f2])).
fof(f1318, plain, (spl102_57 | spl102_41 | spl102_25 | spl102_9), inference(avatar_split_clause, [], [f255, f1059, f1127, f1195, f1263])).
fof(f255, plain, ((e10 = op1(e13, e11)) | (e10 = op1(e12, e11)) | (e10 = op1(e11, e11)) | (e10 = op1(e10, e11))), inference(cnf_transformation, [], [f2])).
fof(f1317, plain, (spl102_46 | spl102_42 | spl102_38 | spl102_34), inference(avatar_split_clause, [], [f256, f1165, f1182, f1199, f1216])).
fof(f256, plain, ((e11 = op1(e11, e13)) | (e11 = op1(e11, e12)) | (e11 = op1(e11, e11)) | (e11 = op1(e11, e10))), inference(cnf_transformation, [], [f2])).
fof(f1316, plain, (spl102_58 | spl102_42 | spl102_26 | spl102_10), inference(avatar_split_clause, [], [f257, f1063, f1131, f1199, f1267])).
fof(f257, plain, ((e11 = op1(e13, e11)) | (e11 = op1(e12, e11)) | (e11 = op1(e11, e11)) | (e11 = op1(e10, e11))), inference(cnf_transformation, [], [f2])).
fof(f1315, plain, (spl102_47 | spl102_43 | spl102_39 | spl102_35), inference(avatar_split_clause, [], [f258, f1169, f1186, f1203, f1220])).
fof(f258, plain, ((e12 = op1(e11, e13)) | (e12 = op1(e11, e12)) | (e12 = op1(e11, e11)) | (e12 = op1(e11, e10))), inference(cnf_transformation, [], [f2])).
fof(f1311, plain, (spl102_29 | spl102_25 | spl102_21 | spl102_17), inference(avatar_split_clause, [], [f262, f1093, f1110, f1127, f1144])).
fof(f262, plain, ((e10 = op1(e12, e13)) | (e10 = op1(e12, e12)) | (e10 = op1(e12, e11)) | (e10 = op1(e12, e10))), inference(cnf_transformation, [], [f2])).
fof(f1310, plain, (spl102_53 | spl102_37 | spl102_21 | spl102_5), inference(avatar_split_clause, [], [f263, f1042, f1110, f1178, f1246])).
fof(f263, plain, ((e10 = op1(e13, e12)) | (e10 = op1(e12, e12)) | (e10 = op1(e11, e12)) | (e10 = op1(e10, e12))), inference(cnf_transformation, [], [f2])).
fof(f1309, plain, (spl102_30 | spl102_26 | spl102_22 | spl102_18), inference(avatar_split_clause, [], [f264, f1097, f1114, f1131, f1148])).
fof(f264, plain, ((e11 = op1(e12, e13)) | (e11 = op1(e12, e12)) | (e11 = op1(e12, e11)) | (e11 = op1(e12, e10))), inference(cnf_transformation, [], [f2])).
fof(f1306, plain, (spl102_55 | spl102_39 | spl102_23 | spl102_7), inference(avatar_split_clause, [], [f267, f1050, f1118, f1186, f1254])).
fof(f267, plain, ((e12 = op1(e13, e12)) | (e12 = op1(e12, e12)) | (e12 = op1(e11, e12)) | (e12 = op1(e10, e12))), inference(cnf_transformation, [], [f2])).
fof(f1305, plain, (spl102_32 | spl102_28 | spl102_24 | spl102_20), inference(avatar_split_clause, [], [f268, f1105, f1122, f1139, f1156])).
fof(f268, plain, ((e13 = op1(e12, e13)) | (e13 = op1(e12, e12)) | (e13 = op1(e12, e11)) | (e13 = op1(e12, e10))), inference(cnf_transformation, [], [f2])).
fof(f1304, plain, (spl102_56 | spl102_40 | spl102_24 | spl102_8), inference(avatar_split_clause, [], [f269, f1054, f1122, f1190, f1258])).
fof(f269, plain, ((e13 = op1(e13, e12)) | (e13 = op1(e12, e12)) | (e13 = op1(e11, e12)) | (e13 = op1(e10, e12))), inference(cnf_transformation, [], [f2])).
fof(f1303, plain, (spl102_13 | spl102_9 | spl102_5 | spl102_1), inference(avatar_split_clause, [], [f270, f1025, f1042, f1059, f1076])).
fof(f270, plain, ((e10 = op1(e13, e13)) | (e10 = op1(e13, e12)) | (e10 = op1(e13, e11)) | (e10 = op1(e13, e10))), inference(cnf_transformation, [], [f2])).
fof(f1302, plain, (spl102_49 | spl102_33 | spl102_17 | spl102_1), inference(avatar_split_clause, [], [f271, f1025, f1093, f1161, f1229])).
fof(f271, plain, ((e10 = op1(e13, e13)) | (e10 = op1(e12, e13)) | (e10 = op1(e11, e13)) | (e10 = op1(e10, e13))), inference(cnf_transformation, [], [f2])).
fof(f1301, plain, (spl102_14 | spl102_10 | spl102_6 | spl102_2), inference(avatar_split_clause, [], [f272, f1029, f1046, f1063, f1080])).
fof(f272, plain, ((e11 = op1(e13, e13)) | (e11 = op1(e13, e12)) | (e11 = op1(e13, e11)) | (e11 = op1(e13, e10))), inference(cnf_transformation, [], [f2])).
fof(f1300, plain, (spl102_50 | spl102_34 | spl102_18 | spl102_2), inference(avatar_split_clause, [], [f273, f1029, f1097, f1165, f1233])).
fof(f273, plain, ((e11 = op1(e13, e13)) | (e11 = op1(e12, e13)) | (e11 = op1(e11, e13)) | (e11 = op1(e10, e13))), inference(cnf_transformation, [], [f2])).
fof(f1298, plain, (spl102_51 | spl102_35 | spl102_19 | spl102_3), inference(avatar_split_clause, [], [f275, f1033, f1101, f1169, f1237])).
fof(f275, plain, ((e12 = op1(e13, e13)) | (e12 = op1(e12, e13)) | (e12 = op1(e11, e13)) | (e12 = op1(e10, e13))), inference(cnf_transformation, [], [f2])).
fof(f1297, plain, (spl102_16 | spl102_12 | spl102_8 | spl102_4), inference(avatar_split_clause, [], [f276, f1037, f1054, f1071, f1088])).
fof(f276, plain, ((e13 = op1(e13, e13)) | (e13 = op1(e13, e12)) | (e13 = op1(e13, e11)) | (e13 = op1(e13, e10))), inference(cnf_transformation, [], [f2])).
fof(f1261, plain, (spl102_53 | spl102_54 | spl102_55 | spl102_56), inference(avatar_split_clause, [], [f232, f1258, f1254, f1250, f1246])).
fof(f232, plain, ((e13 = op1(e10, e12)) | (e12 = op1(e10, e12)) | (e11 = op1(e10, e12)) | (e10 = op1(e10, e12))), inference(cnf_transformation, [], [f1])).
fof(f1, plain, (((e13 = op1(e13, e13)) | (e12 = op1(e13, e13)) | (e11 = op1(e13, e13)) | (e10 = op1(e13, e13))) & ((e13 = op1(e13, e12)) | (e12 = op1(e13, e12)) | (e11 = op1(e13, e12)) | (e10 = op1(e13, e12))) & ((e13 = op1(e13, e11)) | (e12 = op1(e13, e11)) | (e11 = op1(e13, e11)) | (e10 = op1(e13, e11))) & ((e13 = op1(e13, e10)) | (e12 = op1(e13, e10)) | (e11 = op1(e13, e10)) | (e10 = op1(e13, e10))) & ((e13 = op1(e12, e13)) | (e12 = op1(e12, e13)) | (e11 = op1(e12, e13)) | (e10 = op1(e12, e13))) & ((e13 = op1(e12, e12)) | (e12 = op1(e12, e12)) | (e11 = op1(e12, e12)) | (e10 = op1(e12, e12))) & ((e13 = op1(e12, e11)) | (e12 = op1(e12, e11)) | (e11 = op1(e12, e11)) | (e10 = op1(e12, e11))) & ((e13 = op1(e12, e10)) | (e12 = op1(e12, e10)) | (e11 = op1(e12, e10)) | (e10 = op1(e12, e10))) & ((e13 = op1(e11, e13)) | (e12 = op1(e11, e13)) | (e11 = op1(e11, e13)) | (e10 = op1(e11, e13))) & ((e13 = op1(e11, e12)) | (e12 = op1(e11, e12)) | (e11 = op1(e11, e12)) | (e10 = op1(e11, e12))) & ((e13 = op1(e11, e11)) | (e12 = op1(e11, e11)) | (e11 = op1(e11, e11)) | (e10 = op1(e11, e11))) & ((e13 = op1(e11, e10)) | (e12 = op1(e11, e10)) | (e11 = op1(e11, e10)) | (e10 = op1(e11, e10))) & ((e13 = op1(e10, e13)) | (e12 = op1(e10, e13)) | (e11 = op1(e10, e13)) | (e10 = op1(e10, e13))) & ((e13 = op1(e10, e12)) | (e12 = op1(e10, e12)) | (e11 = op1(e10, e12)) | (e10 = op1(e10, e12))) & ((e13 = op1(e10, e11)) | (e12 = op1(e10, e11)) | (e11 = op1(e10, e11)) | (e10 = op1(e10, e11))) & ((op1(e10, e10) = e13) | (op1(e10, e10) = e12) | (op1(e10, e10) = e11) | (e10 = op1(e10, e10)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG109+1.p', ax1)).
fof(f1227, plain, (spl102_45 | spl102_46 | spl102_47 | spl102_48), inference(avatar_split_clause, [], [f234, f1224, f1220, f1216, f1212])).
fof(f234, plain, ((e13 = op1(e11, e10)) | (e12 = op1(e11, e10)) | (e11 = op1(e11, e10)) | (e10 = op1(e11, e10))), inference(cnf_transformation, [], [f1])).
fof(f1210, plain, (spl102_41 | spl102_42 | spl102_43 | spl102_44), inference(avatar_split_clause, [], [f235, f1207, f1203, f1199, f1195])).
fof(f235, plain, ((e13 = op1(e11, e11)) | (e12 = op1(e11, e11)) | (e11 = op1(e11, e11)) | (e10 = op1(e11, e11))), inference(cnf_transformation, [], [f1])).
fof(f1193, plain, (spl102_37 | spl102_38 | spl102_39 | spl102_40), inference(avatar_split_clause, [], [f236, f1190, f1186, f1182, f1178])).
fof(f236, plain, ((e13 = op1(e11, e12)) | (e12 = op1(e11, e12)) | (e11 = op1(e11, e12)) | (e10 = op1(e11, e12))), inference(cnf_transformation, [], [f1])).
fof(f1176, plain, (spl102_33 | spl102_34 | spl102_35 | spl102_36), inference(avatar_split_clause, [], [f237, f1173, f1169, f1165, f1161])).
fof(f237, plain, ((e13 = op1(e11, e13)) | (e12 = op1(e11, e13)) | (e11 = op1(e11, e13)) | (e10 = op1(e11, e13))), inference(cnf_transformation, [], [f1])).
fof(f1159, plain, (spl102_29 | spl102_30 | spl102_31 | spl102_32), inference(avatar_split_clause, [], [f238, f1156, f1152, f1148, f1144])).
fof(f238, plain, ((e13 = op1(e12, e10)) | (e12 = op1(e12, e10)) | (e11 = op1(e12, e10)) | (e10 = op1(e12, e10))), inference(cnf_transformation, [], [f1])).
fof(f1142, plain, (spl102_25 | spl102_26 | spl102_27 | spl102_28), inference(avatar_split_clause, [], [f239, f1139, f1135, f1131, f1127])).
fof(f239, plain, ((e13 = op1(e12, e11)) | (e12 = op1(e12, e11)) | (e11 = op1(e12, e11)) | (e10 = op1(e12, e11))), inference(cnf_transformation, [], [f1])).
fof(f1125, plain, (spl102_21 | spl102_22 | spl102_23 | spl102_24), inference(avatar_split_clause, [], [f240, f1122, f1118, f1114, f1110])).
fof(f240, plain, ((e13 = op1(e12, e12)) | (e12 = op1(e12, e12)) | (e11 = op1(e12, e12)) | (e10 = op1(e12, e12))), inference(cnf_transformation, [], [f1])).
fof(f1108, plain, (spl102_17 | spl102_18 | spl102_19 | spl102_20), inference(avatar_split_clause, [], [f241, f1105, f1101, f1097, f1093])).
fof(f241, plain, ((e13 = op1(e12, e13)) | (e12 = op1(e12, e13)) | (e11 = op1(e12, e13)) | (e10 = op1(e12, e13))), inference(cnf_transformation, [], [f1])).
fof(f1074, plain, (spl102_9 | spl102_10 | spl102_11 | spl102_12), inference(avatar_split_clause, [], [f243, f1071, f1067, f1063, f1059])).
fof(f243, plain, ((e13 = op1(e13, e11)) | (e12 = op1(e13, e11)) | (e11 = op1(e13, e11)) | (e10 = op1(e13, e11))), inference(cnf_transformation, [], [f1])).
fof(f1057, plain, (spl102_5 | spl102_6 | spl102_7 | spl102_8), inference(avatar_split_clause, [], [f244, f1054, f1050, f1046, f1042])).
fof(f244, plain, ((e13 = op1(e13, e12)) | (e12 = op1(e13, e12)) | (e11 = op1(e13, e12)) | (e10 = op1(e13, e12))), inference(cnf_transformation, [], [f1])).
fof(f1040, plain, (spl102_1 | spl102_2 | spl102_3 | spl102_4), inference(avatar_split_clause, [], [f245, f1037, f1033, f1029, f1025])).
fof(f245, plain, ((e13 = op1(e13, e13)) | (e12 = op1(e13, e13)) | (e11 = op1(e13, e13)) | (e10 = op1(e13, e13))), inference(cnf_transformation, [], [f1])).