fof(f4099, plain, $false, inference(avatar_sat_refutation, [], [f747, f764, f798, f815, f832, f866, f917, f934, f968, f988, f993, f996, f999, f1002, f1004, f1005, f1016, f1017, f1051, f1102, f1119, f1153, f1170, f1221, f1238, f1272, f1294, f1295, f1296, f1298, f1300, f1301, f1306, f1308, f1321, f1330, f1331, f1339, f1351, f1360, f1367, f1371, f1379, f1386, f1389, f1390, f1401, f1411, f1419, f1426, f1429, f1430, f1440, f1446, f1447, f1450, f1451, f1461, f1466, f1469, f1471, f1493, f1504, f1530, f1553, f1558, f1561, f1587, f1590, f1591, f1599, f1611, f1620, f1627, f1631, f1639, f1651, f1661, f1666, f1667, f1670, f1679, f1686, f1690, f1700, f1706, f1707, f1711, f1721, f1731, f1748, f1764, f1790, f1813, f1818, f1821, f1843, f1845, f2021, f2407, f2559, f2580, f2581, f2588, f2607, f2638, f2639, f2656, f2664, f2674, f2678, f2712, f2713, f2715, f2732, f2733, f2734, f2761, f2762, f2764, f2765, f2831, f2832, f2833, f2884, f2897, f2899, f2913, f2943, f2953, f2965, f2967, f2969, f2974, f2979, f2996, f3003, f3012, f3015, f3028, f3032, f3042, f3070, f3094, f3112, f3134, f3137, f3139, f3143, f3167, f3198, f3223, f3239, f3262, f3266, f3291, f3304, f3307, f3308, f3335, f3357, f3364, f3370, f3391, f3398, f3399, f3403, f3448, f3451, f3453, f3471, f3524, f3532, f3552, f3561, f3569, f3582, f3593, f3626, f3632, f3647, f3668, f3680, f3697, f3712, f3730, f3743, f3747, f3753, f3763, f3767, f3776, f3782, f3783, f3790, f3800, f3864, f3872, f3876, f3882, f3895, f3922, f3923, f3924, f3936, f3959, f3972, f4027, f4046, f4047, f4075, f4093])).
fof(f4093, plain, (~ spl48_8 | ~ spl48_12), inference(avatar_split_clause, [], [f4090, f761, f744])).
fof(f744, plain, (spl48_8 <=> (e13 = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl48_8])])).
fof(f761, plain, (spl48_12 <=> (e13 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl48_12])])).
fof(f4090, plain, (~ (e13 = op1(e13, e12)) | ~ spl48_12), inference(backward_demodulation, [], [f261, f763])).
fof(f763, plain, ((e13 = op1(e13, e11)) | ~ spl48_12), inference(avatar_component_clause, [], [f761])).
fof(f261, plain, ~ (op1(e13, e11) = op1(e13, e12)), inference(cnf_transformation, [], [f5])).
fof(f5, plain, (~ (op1(e13, e12) = op1(e13, e13)) & ~ (op1(e13, e11) = op1(e13, e13)) & ~ (op1(e13, e11) = op1(e13, e12)) & ~ (op1(e13, e10) = op1(e13, e13)) & ~ (op1(e13, e10) = op1(e13, e12)) & ~ (op1(e13, e10) = op1(e13, e11)) & ~ (op1(e12, e12) = op1(e12, e13)) & ~ (op1(e12, e11) = op1(e12, e13)) & ~ (op1(e12, e11) = op1(e12, e12)) & ~ (op1(e12, e10) = op1(e12, e13)) & ~ (op1(e12, e10) = op1(e12, e12)) & ~ (op1(e12, e10) = op1(e12, e11)) & ~ (op1(e11, e12) = op1(e11, e13)) & ~ (op1(e11, e11) = op1(e11, e13)) & ~ (op1(e11, e11) = op1(e11, e12)) & ~ (op1(e11, e10) = op1(e11, e13)) & ~ (op1(e11, e10) = op1(e11, e12)) & ~ (op1(e11, e10) = op1(e11, e11)) & ~ (op1(e10, e12) = op1(e10, e13)) & ~ (op1(e10, e11) = op1(e10, e13)) & ~ (op1(e10, e11) = op1(e10, e12)) & ~ (op1(e10, e10) = op1(e10, e13)) & ~ (op1(e10, e10) = op1(e10, e12)) & ~ (op1(e10, e10) = op1(e10, e11)) & ~ (op1(e12, e13) = op1(e13, e13)) & ~ (op1(e11, e13) = op1(e13, e13)) & ~ (op1(e11, e13) = op1(e12, e13)) & ~ (op1(e10, e13) = op1(e13, e13)) & ~ (op1(e10, e13) = op1(e12, e13)) & ~ (op1(e10, e13) = op1(e11, e13)) & ~ (op1(e12, e12) = op1(e13, e12)) & ~ (op1(e11, e12) = op1(e13, e12)) & ~ (op1(e11, e12) = op1(e12, e12)) & ~ (op1(e10, e12) = op1(e13, e12)) & ~ (op1(e10, e12) = op1(e12, e12)) & ~ (op1(e10, e12) = op1(e11, e12)) & ~ (op1(e12, e11) = op1(e13, e11)) & ~ (op1(e11, e11) = op1(e13, e11)) & ~ (op1(e11, e11) = op1(e12, e11)) & ~ (op1(e10, e11) = op1(e13, e11)) & ~ (op1(e10, e11) = op1(e12, e11)) & ~ (op1(e10, e11) = op1(e11, e11)) & ~ (op1(e12, e10) = op1(e13, e10)) & ~ (op1(e11, e10) = op1(e13, e10)) & ~ (op1(e11, e10) = op1(e12, e10)) & ~ (op1(e10, e10) = op1(e13, e10)) & ~ (op1(e10, e10) = op1(e12, e10)) & ~ (op1(e10, e10) = op1(e11, e10))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG126+1.p', ax5)).
fof(f4075, plain, (~ spl48_8 | ~ spl48_24), inference(avatar_split_clause, [], [f4070, f812, f744])).
fof(f812, plain, (spl48_24 <=> (e13 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl48_24])])).
fof(f4070, plain, (~ (e13 = op1(e13, e12)) | ~ spl48_24), inference(backward_demodulation, [], [f233, f814])).
fof(f814, plain, ((e13 = op1(e12, e12)) | ~ spl48_24), inference(avatar_component_clause, [], [f812])).
fof(f233, plain, ~ (op1(e12, e12) = op1(e13, e12)), inference(cnf_transformation, [], [f5])).
fof(f4047, plain, (~ spl48_10 | ~ spl48_58), inference(avatar_split_clause, [], [f4044, f957, f753])).
fof(f753, plain, (spl48_10 <=> (e11 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl48_10])])).
fof(f957, plain, (spl48_58 <=> (e11 = op1(e10, e11))), introduced(avatar_definition, [new_symbols(naming, [spl48_58])])).
fof(f4044, plain, (~ (e11 = op1(e13, e11)) | ~ spl48_58), inference(backward_demodulation, [], [f224, f959])).
fof(f959, plain, ((e11 = op1(e10, e11)) | ~ spl48_58), inference(avatar_component_clause, [], [f957])).
fof(f224, plain, ~ (op1(e10, e11) = op1(e13, e11)), inference(cnf_transformation, [], [f5])).
fof(f4046, plain, (~ spl48_50 | ~ spl48_58), inference(avatar_split_clause, [], [f4043, f957, f923])).
fof(f923, plain, (spl48_50 <=> (e11 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl48_50])])).
fof(f4043, plain, (~ (e11 = op1(e10, e13)) | ~ spl48_58), inference(backward_demodulation, [], [f244, f959])).
fof(f244, plain, ~ (op1(e10, e11) = op1(e10, e13)), inference(cnf_transformation, [], [f5])).
fof(f4027, plain, (~ spl48_86 | spl48_91 | ~ spl48_180), inference(avatar_contradiction_clause, [], [f4026])).
fof(f4026, plain, ($false | (~ spl48_86 | spl48_91 | ~ spl48_180)), inference(subsumption_resolution, [], [f4024, f1130])).
fof(f1130, plain, (~ (e22 = op2(e22, e21)) | spl48_91), inference(avatar_component_clause, [], [f1129])).
fof(f1129, plain, (spl48_91 <=> (e22 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl48_91])])).
fof(f4024, plain, ((e22 = op2(e22, e21)) | (~ spl48_86 | ~ spl48_180)), inference(backward_demodulation, [], [f1747, f1110])).
fof(f1110, plain, ((e21 = op2(e22, e22)) | ~ spl48_86), inference(avatar_component_clause, [], [f1108])).
fof(f1108, plain, (spl48_86 <=> (e21 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl48_86])])).
fof(f1747, plain, ((e22 = op2(e22, op2(e22, e22))) | ~ spl48_180), inference(avatar_component_clause, [], [f1745])).
fof(f1745, plain, (spl48_180 <=> (e22 = op2(e22, op2(e22, e22)))), introduced(avatar_definition, [new_symbols(naming, [spl48_180])])).
fof(f3972, plain, (~ spl48_71 | ~ spl48_79), inference(avatar_split_clause, [], [f3971, f1078, f1044])).
fof(f1044, plain, (spl48_71 <=> (e22 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl48_71])])).
fof(f1078, plain, (spl48_79 <=> (e22 = op2(e23, e20))), introduced(avatar_definition, [new_symbols(naming, [spl48_79])])).
fof(f3971, plain, (~ (e22 = op2(e23, e22)) | ~ spl48_79), inference(forward_demodulation, [], [f307, f1080])).
fof(f1080, plain, ((e22 = op2(e23, e20)) | ~ spl48_79), inference(avatar_component_clause, [], [f1078])).
fof(f307, plain, ~ (op2(e23, e20) = op2(e23, e22)), inference(cnf_transformation, [], [f6])).
fof(f6, plain, (~ (op2(e23, e22) = op2(e23, e23)) & ~ (op2(e23, e21) = op2(e23, e23)) & ~ (op2(e23, e21) = op2(e23, e22)) & ~ (op2(e23, e20) = op2(e23, e23)) & ~ (op2(e23, e20) = op2(e23, e22)) & ~ (op2(e23, e20) = op2(e23, e21)) & ~ (op2(e22, e22) = op2(e22, e23)) & ~ (op2(e22, e21) = op2(e22, e23)) & ~ (op2(e22, e21) = op2(e22, e22)) & ~ (op2(e22, e20) = op2(e22, e23)) & ~ (op2(e22, e20) = op2(e22, e22)) & ~ (op2(e22, e20) = op2(e22, e21)) & ~ (op2(e21, e22) = op2(e21, e23)) & ~ (op2(e21, e21) = op2(e21, e23)) & ~ (op2(e21, e21) = op2(e21, e22)) & ~ (op2(e21, e20) = op2(e21, e23)) & ~ (op2(e21, e20) = op2(e21, e22)) & ~ (op2(e21, e20) = op2(e21, e21)) & ~ (op2(e20, e22) = op2(e20, e23)) & ~ (op2(e20, e21) = op2(e20, e23)) & ~ (op2(e20, e21) = op2(e20, e22)) & ~ (op2(e20, e20) = op2(e20, e23)) & ~ (op2(e20, e20) = op2(e20, e22)) & ~ (op2(e20, e20) = op2(e20, e21)) & ~ (op2(e22, e23) = op2(e23, e23)) & ~ (op2(e21, e23) = op2(e23, e23)) & ~ (op2(e21, e23) = op2(e22, e23)) & ~ (op2(e20, e23) = op2(e23, e23)) & ~ (op2(e20, e23) = op2(e22, e23)) & ~ (op2(e20, e23) = op2(e21, e23)) & ~ (op2(e22, e22) = op2(e23, e22)) & ~ (op2(e21, e22) = op2(e23, e22)) & ~ (op2(e21, e22) = op2(e22, e22)) & ~ (op2(e20, e22) = op2(e23, e22)) & ~ (op2(e20, e22) = op2(e22, e22)) & ~ (op2(e20, e22) = op2(e21, e22)) & ~ (op2(e22, e21) = op2(e23, e21)) & ~ (op2(e21, e21) = op2(e23, e21)) & ~ (op2(e21, e21) = op2(e22, e21)) & ~ (op2(e20, e21) = op2(e23, e21)) & ~ (op2(e20, e21) = op2(e22, e21)) & ~ (op2(e20, e21) = op2(e21, e21)) & ~ (op2(e22, e20) = op2(e23, e20)) & ~ (op2(e21, e20) = op2(e23, e20)) & ~ (op2(e21, e20) = op2(e22, e20)) & ~ (op2(e20, e20) = op2(e23, e20)) & ~ (op2(e20, e20) = op2(e22, e20)) & ~ (op2(e20, e20) = op2(e21, e20))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG126+1.p', ax6)).
fof(f3959, plain, (~ spl48_6 | ~ spl48_10), inference(avatar_split_clause, [], [f3954, f753, f736])).
fof(f736, plain, (spl48_6 <=> (e11 = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl48_6])])).
fof(f3954, plain, (~ (e11 = op1(e13, e12)) | ~ spl48_10), inference(backward_demodulation, [], [f261, f755])).
fof(f755, plain, ((e11 = op1(e13, e11)) | ~ spl48_10), inference(avatar_component_clause, [], [f753])).
fof(f3936, plain, (~ spl48_6 | ~ spl48_22), inference(avatar_split_clause, [], [f3929, f804, f736])).
fof(f804, plain, (spl48_22 <=> (e11 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl48_22])])).
fof(f3929, plain, (~ (e11 = op1(e13, e12)) | ~ spl48_22), inference(backward_demodulation, [], [f233, f806])).
fof(f806, plain, ((e11 = op1(e12, e12)) | ~ spl48_22), inference(avatar_component_clause, [], [f804])).
fof(f3924, plain, (spl48_18 | ~ spl48_28 | ~ spl48_147), inference(avatar_split_clause, [], [f3919, f1490, f829, f787])).
fof(f787, plain, (spl48_18 <=> (e11 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl48_18])])).
fof(f829, plain, (spl48_28 <=> (e13 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl48_28])])).
fof(f1490, plain, (spl48_147 <=> (e11 = op1(e12, op1(e12, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl48_147])])).
fof(f3919, plain, ((e11 = op1(e12, e13)) | (~ spl48_28 | ~ spl48_147)), inference(backward_demodulation, [], [f1492, f831])).
fof(f831, plain, ((e13 = op1(e12, e11)) | ~ spl48_28), inference(avatar_component_clause, [], [f829])).
fof(f1492, plain, ((e11 = op1(e12, op1(e12, e11))) | ~ spl48_147), inference(avatar_component_clause, [], [f1490])).
fof(f3923, plain, (~ spl48_20 | ~ spl48_28), inference(avatar_split_clause, [], [f3918, f829, f795])).
fof(f795, plain, (spl48_20 <=> (e13 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl48_20])])).
fof(f3918, plain, (~ (e13 = op1(e12, e13)) | ~ spl48_28), inference(backward_demodulation, [], [f256, f831])).
fof(f256, plain, ~ (op1(e12, e11) = op1(e12, e13)), inference(cnf_transformation, [], [f5])).
fof(f3922, plain, (~ spl48_12 | ~ spl48_28), inference(avatar_split_clause, [], [f3916, f829, f761])).
fof(f3916, plain, (~ (e13 = op1(e13, e11)) | ~ spl48_28), inference(backward_demodulation, [], [f227, f831])).
fof(f227, plain, ~ (op1(e12, e11) = op1(e13, e11)), inference(cnf_transformation, [], [f5])).
fof(f3895, plain, (~ spl48_18 | ~ spl48_50), inference(avatar_split_clause, [], [f3888, f923, f787])).
fof(f3888, plain, (~ (e11 = op1(e12, e13)) | ~ spl48_50), inference(backward_demodulation, [], [f235, f925])).
fof(f925, plain, ((e11 = op1(e10, e13)) | ~ spl48_50), inference(avatar_component_clause, [], [f923])).
fof(f235, plain, ~ (op1(e10, e13) = op1(e12, e13)), inference(cnf_transformation, [], [f5])).
fof(f3882, plain, (~ spl48_51 | ~ spl48_59), inference(avatar_split_clause, [], [f3878, f961, f927])).
fof(f927, plain, (spl48_51 <=> (e12 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl48_51])])).
fof(f961, plain, (spl48_59 <=> (e12 = op1(e10, e11))), introduced(avatar_definition, [new_symbols(naming, [spl48_59])])).
fof(f3878, plain, (~ (e12 = op1(e10, e13)) | ~ spl48_59), inference(backward_demodulation, [], [f244, f963])).
fof(f963, plain, ((e12 = op1(e10, e11)) | ~ spl48_59), inference(avatar_component_clause, [], [f961])).
fof(f3876, plain, (~ spl48_87 | spl48_212), inference(avatar_contradiction_clause, [], [f3875])).
fof(f3875, plain, ($false | (~ spl48_87 | spl48_212)), inference(subsumption_resolution, [], [f3874, f1963])).
fof(f1963, plain, (~ (e22 = h3(e10)) | spl48_212), inference(avatar_component_clause, [], [f1961])).
fof(f1961, plain, (spl48_212 <=> (e22 = h3(e10))), introduced(avatar_definition, [new_symbols(naming, [spl48_212])])).
fof(f3874, plain, ((e22 = h3(e10)) | ~ spl48_87), inference(backward_demodulation, [], [f611, f1114])).
fof(f1114, plain, ((e22 = op2(e22, e22)) | ~ spl48_87), inference(avatar_component_clause, [], [f1112])).
fof(f1112, plain, (spl48_87 <=> (e22 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl48_87])])).
fof(f611, plain, (op2(e22, e22) = h3(e10)), inference(cnf_transformation, [], [f16])).
fof(f16, plain, ((op2(op2(e22, e22), op2(e22, e22)) = h3(e13)) & (h3(e12) = op2(op2(op2(e22, e22), op2(e22, e22)), op2(e22, e22))) & (op2(e22, e22) = h3(e10)) & (e22 = h3(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG126+1.p', ax16)).
fof(f3872, plain, (~ spl48_91 | ~ spl48_92), inference(avatar_contradiction_clause, [], [f3871])).
fof(f3871, plain, ($false | (~ spl48_91 | ~ spl48_92)), inference(subsumption_resolution, [], [f3870, f323])).
fof(f323, plain, ~ (e22 = e23), inference(cnf_transformation, [], [f8])).
fof(f8, plain, (~ (e22 = e23) & ~ (e21 = e23) & ~ (e21 = e22) & ~ (e20 = e23) & ~ (e20 = e22) & ~ (e20 = e21)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG126+1.p', ax8)).
fof(f3870, plain, ((e22 = e23) | (~ spl48_91 | ~ spl48_92)), inference(forward_demodulation, [], [f1135, f1131])).
fof(f1131, plain, ((e22 = op2(e22, e21)) | ~ spl48_91), inference(avatar_component_clause, [], [f1129])).
fof(f1135, plain, ((e23 = op2(e22, e21)) | ~ spl48_92), inference(avatar_component_clause, [], [f1133])).
fof(f1133, plain, (spl48_92 <=> (e23 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl48_92])])).
fof(f3864, plain, (~ spl48_105 | ~ spl48_107), inference(avatar_contradiction_clause, [], [f3863])).
fof(f3863, plain, ($false | (~ spl48_105 | ~ spl48_107)), inference(subsumption_resolution, [], [f3862, f319])).
fof(f319, plain, ~ (e20 = e22), inference(cnf_transformation, [], [f8])).
fof(f3862, plain, ((e20 = e22) | (~ spl48_105 | ~ spl48_107)), inference(forward_demodulation, [], [f1199, f1191])).
fof(f1191, plain, ((e20 = op2(e21, e21)) | ~ spl48_105), inference(avatar_component_clause, [], [f1189])).
fof(f1189, plain, (spl48_105 <=> (e20 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl48_105])])).
fof(f1199, plain, ((e22 = op2(e21, e21)) | ~ spl48_107), inference(avatar_component_clause, [], [f1197])).
fof(f1197, plain, (spl48_107 <=> (e22 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl48_107])])).
fof(f3800, plain, (~ spl48_41 | ~ spl48_43), inference(avatar_contradiction_clause, [], [f3799])).
fof(f3799, plain, ($false | (~ spl48_41 | ~ spl48_43)), inference(subsumption_resolution, [], [f3798, f313])).
fof(f313, plain, ~ (e10 = e12), inference(cnf_transformation, [], [f7])).
fof(f7, plain, (~ (e12 = e13) & ~ (e11 = e13) & ~ (e11 = e12) & ~ (e10 = e13) & ~ (e10 = e12) & ~ (e10 = e11)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG126+1.p', ax7)).
fof(f3798, plain, ((e10 = e12) | (~ spl48_41 | ~ spl48_43)), inference(forward_demodulation, [], [f895, f887])).
fof(f887, plain, ((e10 = op1(e11, e11)) | ~ spl48_41), inference(avatar_component_clause, [], [f885])).
fof(f885, plain, (spl48_41 <=> (e10 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl48_41])])).
fof(f895, plain, ((e12 = op1(e11, e11)) | ~ spl48_43), inference(avatar_component_clause, [], [f893])).
fof(f893, plain, (spl48_43 <=> (e12 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl48_43])])).
fof(f3790, plain, (~ spl48_34 | ~ spl48_46), inference(avatar_split_clause, [], [f3394, f906, f855])).
fof(f855, plain, (spl48_34 <=> (e11 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl48_34])])).
fof(f906, plain, (spl48_46 <=> (e11 = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl48_46])])).
fof(f3394, plain, (~ (e11 = op1(e11, e13)) | ~ spl48_46), inference(backward_demodulation, [], [f248, f908])).
fof(f908, plain, ((e11 = op1(e11, e10)) | ~ spl48_46), inference(avatar_component_clause, [], [f906])).
fof(f248, plain, ~ (op1(e11, e10) = op1(e11, e13)), inference(cnf_transformation, [], [f5])).
fof(f3783, plain, (~ spl48_23 | ~ spl48_27), inference(avatar_split_clause, [], [f3055, f825, f808])).
fof(f808, plain, (spl48_23 <=> (e12 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl48_23])])).
fof(f825, plain, (spl48_27 <=> (e12 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl48_27])])).
fof(f3055, plain, (~ (e12 = op1(e12, e12)) | ~ spl48_27), inference(backward_demodulation, [], [f255, f827])).
fof(f827, plain, ((e12 = op1(e12, e11)) | ~ spl48_27), inference(avatar_component_clause, [], [f825])).
fof(f255, plain, ~ (op1(e12, e11) = op1(e12, e12)), inference(cnf_transformation, [], [f5])).
fof(f3782, plain, (~ spl48_6 | ~ spl48_70 | spl48_274), inference(avatar_contradiction_clause, [], [f3781])).
fof(f3781, plain, ($false | (~ spl48_6 | ~ spl48_70 | spl48_274)), inference(subsumption_resolution, [], [f3780, f1042])).
fof(f1042, plain, ((e21 = op2(e23, e22)) | ~ spl48_70), inference(avatar_component_clause, [], [f1040])).
fof(f1040, plain, (spl48_70 <=> (e21 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl48_70])])).
fof(f3780, plain, (~ (e21 = op2(e23, e22)) | (~ spl48_6 | spl48_274)), inference(forward_demodulation, [], [f3779, f606])).
fof(f606, plain, (e21 = h2(e11)), inference(cnf_transformation, [], [f15])).
fof(f15, plain, ((op2(op2(e21, e21), op2(e21, e21)) = h2(e13)) & (op2(op2(op2(e21, e21), op2(e21, e21)), op2(e21, e21)) = h2(e12)) & (op2(e21, e21) = h2(e10)) & (e21 = h2(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG126+1.p', ax15)).
fof(f3779, plain, (~ (op2(e23, e22) = h2(e11)) | (~ spl48_6 | spl48_274)), inference(forward_demodulation, [], [f2350, f738])).
fof(f738, plain, ((e11 = op1(e13, e12)) | ~ spl48_6), inference(avatar_component_clause, [], [f736])).
fof(f2350, plain, (~ (op2(e23, e22) = h2(op1(e13, e12))) | spl48_274), inference(avatar_component_clause, [], [f2348])).
fof(f2348, plain, (spl48_274 <=> (op2(e23, e22) = h2(op1(e13, e12)))), introduced(avatar_definition, [new_symbols(naming, [spl48_274])])).
fof(f3776, plain, (~ spl48_36 | ~ spl48_100 | spl48_281), inference(avatar_contradiction_clause, [], [f3775])).
fof(f3775, plain, ($false | (~ spl48_36 | ~ spl48_100 | spl48_281)), inference(subsumption_resolution, [], [f3774, f1169])).
fof(f1169, plain, ((e23 = op2(e21, e23)) | ~ spl48_100), inference(avatar_component_clause, [], [f1167])).
fof(f1167, plain, (spl48_100 <=> (e23 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl48_100])])).
fof(f3774, plain, (~ (e23 = op2(e21, e23)) | (~ spl48_36 | spl48_281)), inference(forward_demodulation, [], [f3773, f1855])).
fof(f1855, plain, (e23 = h2(e13)), inference(forward_demodulation, [], [f609, f601])).
fof(f601, plain, (e23 = op2(op2(e21, e21), op2(e21, e21))), inference(cnf_transformation, [], [f13])).
fof(f13, plain, ((e23 = op2(op2(e21, e21), op2(e21, e21))) & (e22 = op2(op2(op2(e21, e21), op2(e21, e21)), op2(e21, e21))) & (e20 = op2(e21, e21))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG126+1.p', ax13)).
fof(f609, plain, (op2(op2(e21, e21), op2(e21, e21)) = h2(e13)), inference(cnf_transformation, [], [f15])).
fof(f3773, plain, (~ (op2(e21, e23) = h2(e13)) | (~ spl48_36 | spl48_281)), inference(forward_demodulation, [], [f2378, f865])).
fof(f865, plain, ((e13 = op1(e11, e13)) | ~ spl48_36), inference(avatar_component_clause, [], [f863])).
fof(f863, plain, (spl48_36 <=> (e13 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl48_36])])).
fof(f2378, plain, (~ (op2(e21, e23) = h2(op1(e11, e13))) | spl48_281), inference(avatar_component_clause, [], [f2376])).
fof(f2376, plain, (spl48_281 <=> (op2(e21, e23) = h2(op1(e11, e13)))), introduced(avatar_definition, [new_symbols(naming, [spl48_281])])).
fof(f3767, plain, (~ spl48_39 | ~ spl48_103 | spl48_282), inference(avatar_contradiction_clause, [], [f3766])).
fof(f3766, plain, ($false | (~ spl48_39 | ~ spl48_103 | spl48_282)), inference(subsumption_resolution, [], [f3765, f1182])).
fof(f1182, plain, ((e22 = op2(e21, e22)) | ~ spl48_103), inference(avatar_component_clause, [], [f1180])).
fof(f1180, plain, (spl48_103 <=> (e22 = op2(e21, e22))), introduced(avatar_definition, [new_symbols(naming, [spl48_103])])).
fof(f3765, plain, (~ (e22 = op2(e21, e22)) | (~ spl48_39 | spl48_282)), inference(forward_demodulation, [], [f3764, f1857])).
fof(f1857, plain, (e22 = h2(e12)), inference(forward_demodulation, [], [f1856, f1844])).
fof(f1844, plain, (e22 = op2(e23, op2(e21, e21))), inference(forward_demodulation, [], [f600, f601])).
fof(f600, plain, (e22 = op2(op2(op2(e21, e21), op2(e21, e21)), op2(e21, e21))), inference(cnf_transformation, [], [f13])).
fof(f1856, plain, (h2(e12) = op2(e23, op2(e21, e21))), inference(forward_demodulation, [], [f608, f601])).
fof(f608, plain, (op2(op2(op2(e21, e21), op2(e21, e21)), op2(e21, e21)) = h2(e12)), inference(cnf_transformation, [], [f15])).
fof(f3764, plain, (~ (op2(e21, e22) = h2(e12)) | (~ spl48_39 | spl48_282)), inference(forward_demodulation, [], [f2382, f878])).
fof(f878, plain, ((e12 = op1(e11, e12)) | ~ spl48_39), inference(avatar_component_clause, [], [f876])).
fof(f876, plain, (spl48_39 <=> (e12 = op1(e11, e12))), introduced(avatar_definition, [new_symbols(naming, [spl48_39])])).
fof(f2382, plain, (~ (op2(e21, e22) = h2(op1(e11, e12))) | spl48_282), inference(avatar_component_clause, [], [f2380])).
fof(f2380, plain, (spl48_282 <=> (op2(e21, e22) = h2(op1(e11, e12)))), introduced(avatar_definition, [new_symbols(naming, [spl48_282])])).
fof(f3763, plain, (~ spl48_29 | ~ spl48_93 | ~ spl48_222 | spl48_280), inference(avatar_contradiction_clause, [], [f3762])).
fof(f3762, plain, ($false | (~ spl48_29 | ~ spl48_93 | ~ spl48_222 | spl48_280)), inference(subsumption_resolution, [], [f3761, f1140])).
fof(f1140, plain, ((e20 = op2(e22, e20)) | ~ spl48_93), inference(avatar_component_clause, [], [f1138])).
fof(f1138, plain, (spl48_93 <=> (e20 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl48_93])])).
fof(f3761, plain, (~ (e20 = op2(e22, e20)) | (~ spl48_29 | ~ spl48_222 | spl48_280)), inference(forward_demodulation, [], [f3760, f2019])).
fof(f2019, plain, ((e20 = h2(e10)) | ~ spl48_222), inference(avatar_component_clause, [], [f2018])).
fof(f2018, plain, (spl48_222 <=> (e20 = h2(e10))), introduced(avatar_definition, [new_symbols(naming, [spl48_222])])).
fof(f3760, plain, (~ (op2(e22, e20) = h2(e10)) | (~ spl48_29 | ~ spl48_222 | spl48_280)), inference(forward_demodulation, [], [f3759, f836])).
fof(f836, plain, ((e10 = op1(e12, e10)) | ~ spl48_29), inference(avatar_component_clause, [], [f834])).
fof(f834, plain, (spl48_29 <=> (e10 = op1(e12, e10))), introduced(avatar_definition, [new_symbols(naming, [spl48_29])])).
fof(f3759, plain, (~ (op2(e22, e20) = h2(op1(e12, e10))) | (~ spl48_222 | spl48_280)), inference(forward_demodulation, [], [f2374, f2019])).
fof(f2374, plain, (~ (h2(op1(e12, e10)) = op2(e22, h2(e10))) | spl48_280), inference(avatar_component_clause, [], [f2372])).
fof(f2372, plain, (spl48_280 <=> (h2(op1(e12, e10)) = op2(e22, h2(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl48_280])])).
fof(f3753, plain, (~ spl48_51 | ~ spl48_115 | ~ spl48_222 | spl48_285), inference(avatar_contradiction_clause, [], [f3752])).
fof(f3752, plain, ($false | (~ spl48_51 | ~ spl48_115 | ~ spl48_222 | spl48_285)), inference(subsumption_resolution, [], [f3751, f1233])).
fof(f1233, plain, ((e22 = op2(e20, e23)) | ~ spl48_115), inference(avatar_component_clause, [], [f1231])).
fof(f1231, plain, (spl48_115 <=> (e22 = op2(e20, e23))), introduced(avatar_definition, [new_symbols(naming, [spl48_115])])).
fof(f3751, plain, (~ (e22 = op2(e20, e23)) | (~ spl48_51 | ~ spl48_222 | spl48_285)), inference(forward_demodulation, [], [f3750, f1857])).
fof(f3750, plain, (~ (op2(e20, e23) = h2(e12)) | (~ spl48_51 | ~ spl48_222 | spl48_285)), inference(forward_demodulation, [], [f3749, f929])).
fof(f929, plain, ((e12 = op1(e10, e13)) | ~ spl48_51), inference(avatar_component_clause, [], [f927])).
fof(f3749, plain, (~ (op2(e20, e23) = h2(op1(e10, e13))) | (~ spl48_222 | spl48_285)), inference(forward_demodulation, [], [f2394, f2019])).
fof(f2394, plain, (~ (h2(op1(e10, e13)) = op2(h2(e10), e23)) | spl48_285), inference(avatar_component_clause, [], [f2392])).
fof(f2392, plain, (spl48_285 <=> (h2(op1(e10, e13)) = op2(h2(e10), e23))), introduced(avatar_definition, [new_symbols(naming, [spl48_285])])).
fof(f3747, plain, (~ spl48_24 | ~ spl48_272 | spl48_278), inference(avatar_contradiction_clause, [], [f3746])).
fof(f3746, plain, ($false | (~ spl48_24 | ~ spl48_272 | spl48_278)), inference(subsumption_resolution, [], [f3745, f1855])).
fof(f3745, plain, (~ (e23 = h2(e13)) | (~ spl48_24 | ~ spl48_272 | spl48_278)), inference(forward_demodulation, [], [f3744, f2311])).
fof(f2311, plain, ((e23 = h3(e10)) | ~ spl48_272), inference(avatar_component_clause, [], [f2310])).
fof(f2310, plain, (spl48_272 <=> (e23 = h3(e10))), introduced(avatar_definition, [new_symbols(naming, [spl48_272])])).
fof(f3744, plain, (~ (h2(e13) = h3(e10)) | (~ spl48_24 | spl48_278)), inference(forward_demodulation, [], [f2366, f814])).
fof(f2366, plain, (~ (h3(e10) = h2(op1(e12, e12))) | spl48_278), inference(avatar_component_clause, [], [f2364])).
fof(f2364, plain, (spl48_278 <=> (h3(e10) = h2(op1(e12, e12)))), introduced(avatar_definition, [new_symbols(naming, [spl48_278])])).
fof(f3743, plain, (~ spl48_15 | spl48_276), inference(avatar_contradiction_clause, [], [f3742])).
fof(f3742, plain, ($false | (~ spl48_15 | spl48_276)), inference(subsumption_resolution, [], [f3741, f1857])).
fof(f3741, plain, (~ (e22 = h2(e12)) | (~ spl48_15 | spl48_276)), inference(forward_demodulation, [], [f2358, f776])).
fof(f776, plain, ((e12 = op1(e13, e10)) | ~ spl48_15), inference(avatar_component_clause, [], [f774])).
fof(f774, plain, (spl48_15 <=> (e12 = op1(e13, e10))), introduced(avatar_definition, [new_symbols(naming, [spl48_15])])).
fof(f2358, plain, (~ (e22 = h2(op1(e13, e10))) | spl48_276), inference(avatar_component_clause, [], [f2356])).
fof(f2356, plain, (spl48_276 <=> (e22 = h2(op1(e13, e10)))), introduced(avatar_definition, [new_symbols(naming, [spl48_276])])).
fof(f3730, plain, (~ spl48_18 | ~ spl48_82 | spl48_277), inference(avatar_contradiction_clause, [], [f3729])).
fof(f3729, plain, ($false | (~ spl48_18 | ~ spl48_82 | spl48_277)), inference(subsumption_resolution, [], [f3728, f1093])).
fof(f1093, plain, ((e21 = op2(e22, e23)) | ~ spl48_82), inference(avatar_component_clause, [], [f1091])).
fof(f1091, plain, (spl48_82 <=> (e21 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl48_82])])).
fof(f3728, plain, (~ (e21 = op2(e22, e23)) | (~ spl48_18 | spl48_277)), inference(forward_demodulation, [], [f3727, f606])).
fof(f3727, plain, (~ (op2(e22, e23) = h2(e11)) | (~ spl48_18 | spl48_277)), inference(forward_demodulation, [], [f2362, f789])).
fof(f789, plain, ((e11 = op1(e12, e13)) | ~ spl48_18), inference(avatar_component_clause, [], [f787])).
fof(f2362, plain, (~ (op2(e22, e23) = h2(op1(e12, e13))) | spl48_277), inference(avatar_component_clause, [], [f2360])).
fof(f2360, plain, (spl48_277 <=> (op2(e22, e23) = h2(op1(e12, e13)))), introduced(avatar_definition, [new_symbols(naming, [spl48_277])])).
fof(f3712, plain, (~ spl48_46 | ~ spl48_110 | ~ spl48_222 | spl48_284), inference(avatar_contradiction_clause, [], [f3711])).
fof(f3711, plain, ($false | (~ spl48_46 | ~ spl48_110 | ~ spl48_222 | spl48_284)), inference(subsumption_resolution, [], [f3710, f1212])).
fof(f1212, plain, ((e21 = op2(e21, e20)) | ~ spl48_110), inference(avatar_component_clause, [], [f1210])).
fof(f1210, plain, (spl48_110 <=> (e21 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl48_110])])).
fof(f3710, plain, (~ (e21 = op2(e21, e20)) | (~ spl48_46 | ~ spl48_222 | spl48_284)), inference(forward_demodulation, [], [f3709, f606])).
fof(f3709, plain, (~ (op2(e21, e20) = h2(e11)) | (~ spl48_46 | ~ spl48_222 | spl48_284)), inference(forward_demodulation, [], [f3708, f908])).
fof(f3708, plain, (~ (op2(e21, e20) = h2(op1(e11, e10))) | (~ spl48_222 | spl48_284)), inference(forward_demodulation, [], [f2390, f2019])).
fof(f2390, plain, (~ (h2(op1(e11, e10)) = op2(e21, h2(e10))) | spl48_284), inference(avatar_component_clause, [], [f2388])).
fof(f2388, plain, (spl48_284 <=> (h2(op1(e11, e10)) = op2(e21, h2(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl48_284])])).
fof(f3697, plain, (~ spl48_12 | ~ spl48_76 | spl48_275), inference(avatar_contradiction_clause, [], [f3696])).
fof(f3696, plain, ($false | (~ spl48_12 | ~ spl48_76 | spl48_275)), inference(subsumption_resolution, [], [f3695, f1067])).
fof(f1067, plain, ((e23 = op2(e23, e21)) | ~ spl48_76), inference(avatar_component_clause, [], [f1065])).
fof(f1065, plain, (spl48_76 <=> (e23 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl48_76])])).
fof(f3695, plain, (~ (e23 = op2(e23, e21)) | (~ spl48_12 | spl48_275)), inference(forward_demodulation, [], [f3694, f1855])).
fof(f3694, plain, (~ (op2(e23, e21) = h2(e13)) | (~ spl48_12 | spl48_275)), inference(forward_demodulation, [], [f2354, f763])).
fof(f2354, plain, (~ (op2(e23, e21) = h2(op1(e13, e11))) | spl48_275), inference(avatar_component_clause, [], [f2352])).
fof(f2352, plain, (spl48_275 <=> (op2(e23, e21) = h2(op1(e13, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl48_275])])).
fof(f3680, plain, (~ spl48_53 | ~ spl48_117 | ~ spl48_222 | spl48_286), inference(avatar_contradiction_clause, [], [f3679])).
fof(f3679, plain, ($false | (~ spl48_53 | ~ spl48_117 | ~ spl48_222 | spl48_286)), inference(subsumption_resolution, [], [f3678, f1242])).
fof(f1242, plain, ((e20 = op2(e20, e22)) | ~ spl48_117), inference(avatar_component_clause, [], [f1240])).
fof(f1240, plain, (spl48_117 <=> (e20 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl48_117])])).
fof(f3678, plain, (~ (e20 = op2(e20, e22)) | (~ spl48_53 | ~ spl48_222 | spl48_286)), inference(forward_demodulation, [], [f3677, f2019])).
fof(f3677, plain, (~ (op2(e20, e22) = h2(e10)) | (~ spl48_53 | ~ spl48_222 | spl48_286)), inference(forward_demodulation, [], [f3676, f938])).
fof(f938, plain, ((e10 = op1(e10, e12)) | ~ spl48_53), inference(avatar_component_clause, [], [f936])).
fof(f936, plain, (spl48_53 <=> (e10 = op1(e10, e12))), introduced(avatar_definition, [new_symbols(naming, [spl48_53])])).
fof(f3676, plain, (~ (op2(e20, e22) = h2(op1(e10, e12))) | (~ spl48_222 | spl48_286)), inference(forward_demodulation, [], [f2398, f2019])).
fof(f2398, plain, (~ (h2(op1(e10, e12)) = op2(h2(e10), e22)) | spl48_286), inference(avatar_component_clause, [], [f2396])).
fof(f2396, plain, (spl48_286 <=> (h2(op1(e10, e12)) = op2(h2(e10), e22))), introduced(avatar_definition, [new_symbols(naming, [spl48_286])])).
fof(f3668, plain, (~ spl48_64 | spl48_288), inference(avatar_contradiction_clause, [], [f3667])).
fof(f3667, plain, ($false | (~ spl48_64 | spl48_288)), inference(subsumption_resolution, [], [f3666, f1855])).
fof(f3666, plain, (~ (e23 = h2(e13)) | (~ spl48_64 | spl48_288)), inference(forward_demodulation, [], [f2406, f984])).
fof(f984, plain, ((op1(e10, e10) = e13) | ~ spl48_64), inference(avatar_component_clause, [], [f982])).
fof(f982, plain, (spl48_64 <=> (op1(e10, e10) = e13)), introduced(avatar_definition, [new_symbols(naming, [spl48_64])])).
fof(f2406, plain, (~ (e23 = h2(op1(e10, e10))) | spl48_288), inference(avatar_component_clause, [], [f2404])).
fof(f2404, plain, (spl48_288 <=> (e23 = h2(op1(e10, e10)))), introduced(avatar_definition, [new_symbols(naming, [spl48_288])])).
fof(f3647, plain, (~ spl48_58 | ~ spl48_122 | ~ spl48_222 | spl48_287), inference(avatar_contradiction_clause, [], [f3646])).
fof(f3646, plain, ($false | (~ spl48_58 | ~ spl48_122 | ~ spl48_222 | spl48_287)), inference(subsumption_resolution, [], [f3645, f1263])).
fof(f1263, plain, ((e21 = op2(e20, e21)) | ~ spl48_122), inference(avatar_component_clause, [], [f1261])).
fof(f1261, plain, (spl48_122 <=> (e21 = op2(e20, e21))), introduced(avatar_definition, [new_symbols(naming, [spl48_122])])).
fof(f3645, plain, (~ (e21 = op2(e20, e21)) | (~ spl48_58 | ~ spl48_222 | spl48_287)), inference(forward_demodulation, [], [f3644, f606])).
fof(f3644, plain, (~ (op2(e20, e21) = h2(e11)) | (~ spl48_58 | ~ spl48_222 | spl48_287)), inference(forward_demodulation, [], [f3643, f959])).
fof(f3643, plain, (~ (op2(e20, e21) = h2(op1(e10, e11))) | (~ spl48_222 | spl48_287)), inference(forward_demodulation, [], [f2402, f2019])).
fof(f2402, plain, (~ (h2(op1(e10, e11)) = op2(h2(e10), e21)) | spl48_287), inference(avatar_component_clause, [], [f2400])).
fof(f2400, plain, (spl48_287 <=> (h2(op1(e10, e11)) = op2(h2(e10), e21))), introduced(avatar_definition, [new_symbols(naming, [spl48_287])])).
fof(f3632, plain, (~ spl48_81 | ~ spl48_93), inference(avatar_split_clause, [], [f3631, f1138, f1087])).
fof(f1087, plain, (spl48_81 <=> (e20 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl48_81])])).
fof(f3631, plain, (~ (e20 = op2(e22, e23)) | ~ spl48_93), inference(backward_demodulation, [], [f302, f1140])).
fof(f302, plain, ~ (op2(e22, e20) = op2(e22, e23)), inference(cnf_transformation, [], [f6])).
fof(f3626, plain, (spl48_225 | ~ spl48_115 | ~ spl48_232 | ~ spl48_308), inference(avatar_split_clause, [], [f3624, f2553, f2067, f1231, f2032])).
fof(f2032, plain, (spl48_225 <=> (e22 = h1(e12))), introduced(avatar_definition, [new_symbols(naming, [spl48_225])])).
fof(f2067, plain, (spl48_232 <=> (e20 = h1(e13))), introduced(avatar_definition, [new_symbols(naming, [spl48_232])])).
fof(f2553, plain, (spl48_308 <=> (e23 = h1(e10))), introduced(avatar_definition, [new_symbols(naming, [spl48_308])])).
fof(f3624, plain, ((e22 = h1(e12)) | (~ spl48_115 | ~ spl48_232 | ~ spl48_308)), inference(backward_demodulation, [], [f3601, f1233])).
fof(f3601, plain, ((op2(e20, e23) = h1(e12)) | (~ spl48_232 | ~ spl48_308)), inference(backward_demodulation, [], [f3165, f2068])).
fof(f2068, plain, ((e20 = h1(e13)) | ~ spl48_232), inference(avatar_component_clause, [], [f2067])).
fof(f3165, plain, ((h1(e12) = op2(h1(e13), e23)) | ~ spl48_308), inference(backward_demodulation, [], [f1854, f2554])).
fof(f2554, plain, ((e23 = h1(e10)) | ~ spl48_308), inference(avatar_component_clause, [], [f2553])).
fof(f1854, plain, (h1(e12) = op2(h1(e13), h1(e10))), inference(backward_demodulation, [], [f1846, f603])).
fof(f603, plain, (op2(e20, e20) = h1(e10)), inference(cnf_transformation, [], [f14])).
fof(f14, plain, ((op2(op2(e20, e20), op2(e20, e20)) = h1(e13)) & (h1(e12) = op2(op2(op2(e20, e20), op2(e20, e20)), op2(e20, e20))) & (op2(e20, e20) = h1(e10)) & (e20 = h1(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG126+1.p', ax14)).
fof(f1846, plain, (h1(e12) = op2(h1(e13), op2(e20, e20))), inference(forward_demodulation, [], [f604, f605])).
fof(f605, plain, (op2(op2(e20, e20), op2(e20, e20)) = h1(e13)), inference(cnf_transformation, [], [f14])).
fof(f604, plain, (h1(e12) = op2(op2(op2(e20, e20), op2(e20, e20)), op2(e20, e20))), inference(cnf_transformation, [], [f14])).
fof(f3593, plain, (~ spl48_41 | spl48_283), inference(avatar_contradiction_clause, [], [f3592])).
fof(f3592, plain, ($false | (~ spl48_41 | spl48_283)), inference(trivial_inequality_removal, [], [f3591])).
fof(f3591, plain, (~ (h2(e10) = h2(e10)) | (~ spl48_41 | spl48_283)), inference(forward_demodulation, [], [f2386, f887])).
fof(f2386, plain, (~ (h2(e10) = h2(op1(e11, e11))) | spl48_283), inference(avatar_component_clause, [], [f2384])).
fof(f2384, plain, (spl48_283 <=> (h2(e10) = h2(op1(e11, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl48_283])])).
fof(f3582, plain, (~ spl48_114 | ~ spl48_122), inference(avatar_split_clause, [], [f3581, f1261, f1227])).
fof(f1227, plain, (spl48_114 <=> (e21 = op2(e20, e23))), introduced(avatar_definition, [new_symbols(naming, [spl48_114])])).
fof(f3581, plain, (~ (e21 = op2(e20, e23)) | ~ spl48_122), inference(forward_demodulation, [], [f292, f1263])).
fof(f292, plain, ~ (op2(e20, e21) = op2(e20, e23)), inference(cnf_transformation, [], [f6])).
fof(f3569, plain, (~ spl48_83 | ~ spl48_91), inference(avatar_split_clause, [], [f3568, f1129, f1095])).
fof(f1095, plain, (spl48_83 <=> (e22 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl48_83])])).
fof(f3568, plain, (~ (e22 = op2(e22, e23)) | ~ spl48_91), inference(forward_demodulation, [], [f304, f1131])).
fof(f304, plain, ~ (op2(e22, e21) = op2(e22, e23)), inference(cnf_transformation, [], [f6])).
fof(f3561, plain, (~ spl48_95 | ~ spl48_79), inference(avatar_split_clause, [], [f3560, f1078, f1146])).
fof(f1146, plain, (spl48_95 <=> (e22 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl48_95])])).
fof(f3560, plain, (~ (e22 = op2(e22, e20)) | ~ spl48_79), inference(forward_demodulation, [], [f269, f1080])).
fof(f269, plain, ~ (op2(e22, e20) = op2(e23, e20)), inference(cnf_transformation, [], [f6])).
fof(f3552, plain, (~ spl48_77 | ~ spl48_79), inference(avatar_contradiction_clause, [], [f3551])).
fof(f3551, plain, ($false | (~ spl48_77 | ~ spl48_79)), inference(subsumption_resolution, [], [f3550, f319])).
fof(f3550, plain, ((e20 = e22) | (~ spl48_77 | ~ spl48_79)), inference(forward_demodulation, [], [f1080, f1072])).
fof(f1072, plain, ((e20 = op2(e23, e20)) | ~ spl48_77), inference(avatar_component_clause, [], [f1070])).
fof(f1070, plain, (spl48_77 <=> (e20 = op2(e23, e20))), introduced(avatar_definition, [new_symbols(naming, [spl48_77])])).
fof(f3532, plain, (~ spl48_225 | ~ spl48_103 | ~ spl48_228 | ~ spl48_308), inference(avatar_split_clause, [], [f3530, f2553, f2047, f1180, f2032])).
fof(f2047, plain, (spl48_228 <=> (e21 = h1(e13))), introduced(avatar_definition, [new_symbols(naming, [spl48_228])])).
fof(f3530, plain, (~ (e22 = h1(e12)) | (~ spl48_103 | ~ spl48_228 | ~ spl48_308)), inference(backward_demodulation, [], [f3506, f1182])).
fof(f3506, plain, (~ (op2(e21, e22) = h1(e12)) | (~ spl48_228 | ~ spl48_308)), inference(backward_demodulation, [], [f299, f3495])).
fof(f3495, plain, ((op2(e21, e23) = h1(e12)) | (~ spl48_228 | ~ spl48_308)), inference(backward_demodulation, [], [f3165, f2048])).
fof(f2048, plain, ((e21 = h1(e13)) | ~ spl48_228), inference(avatar_component_clause, [], [f2047])).
fof(f299, plain, ~ (op2(e21, e22) = op2(e21, e23)), inference(cnf_transformation, [], [f6])).
fof(f3524, plain, (~ spl48_122 | ~ spl48_123), inference(avatar_contradiction_clause, [], [f3523])).
fof(f3523, plain, ($false | (~ spl48_122 | ~ spl48_123)), inference(subsumption_resolution, [], [f3522, f321])).
fof(f321, plain, ~ (e21 = e22), inference(cnf_transformation, [], [f8])).
fof(f3522, plain, ((e21 = e22) | (~ spl48_122 | ~ spl48_123)), inference(backward_demodulation, [], [f1267, f1263])).
fof(f1267, plain, ((e22 = op2(e20, e21)) | ~ spl48_123), inference(avatar_component_clause, [], [f1265])).
fof(f1265, plain, (spl48_123 <=> (e22 = op2(e20, e21))), introduced(avatar_definition, [new_symbols(naming, [spl48_123])])).
fof(f3471, plain, (~ spl48_232 | ~ spl48_1 | ~ spl48_222 | spl48_273 | ~ spl48_308), inference(avatar_split_clause, [], [f3470, f2553, f2344, f2018, f715, f2067])).
fof(f715, plain, (spl48_1 <=> (e10 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl48_1])])).
fof(f2344, plain, (spl48_273 <=> (h4(e10) = h2(op1(e13, e13)))), introduced(avatar_definition, [new_symbols(naming, [spl48_273])])).
fof(f3470, plain, (~ (e20 = h1(e13)) | (~ spl48_1 | ~ spl48_222 | spl48_273 | ~ spl48_308)), inference(forward_demodulation, [], [f3469, f2019])).
fof(f3469, plain, (~ (h1(e13) = h2(e10)) | (~ spl48_1 | spl48_273 | ~ spl48_308)), inference(forward_demodulation, [], [f3468, f3168])).
fof(f3168, plain, ((h1(e13) = h4(e10)) | ~ spl48_308), inference(backward_demodulation, [], [f615, f3164])).
fof(f3164, plain, ((op2(e23, e23) = h1(e13)) | ~ spl48_308), inference(backward_demodulation, [], [f1853, f2554])).
fof(f1853, plain, (h1(e13) = op2(h1(e10), h1(e10))), inference(backward_demodulation, [], [f605, f603])).
fof(f615, plain, (op2(e23, e23) = h4(e10)), inference(cnf_transformation, [], [f17])).
fof(f17, plain, ((op2(op2(e23, e23), op2(e23, e23)) = h4(e13)) & (h4(e12) = op2(op2(op2(e23, e23), op2(e23, e23)), op2(e23, e23))) & (op2(e23, e23) = h4(e10)) & (e23 = h4(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG126+1.p', ax17)).
fof(f3468, plain, (~ (h2(e10) = h4(e10)) | (~ spl48_1 | spl48_273)), inference(forward_demodulation, [], [f2346, f717])).
fof(f717, plain, ((e10 = op1(e13, e13)) | ~ spl48_1), inference(avatar_component_clause, [], [f715])).
fof(f2346, plain, (~ (h4(e10) = h2(op1(e13, e13))) | spl48_273), inference(avatar_component_clause, [], [f2344])).
fof(f3453, plain, (~ spl48_94 | ~ spl48_110), inference(avatar_split_clause, [], [f3452, f1210, f1142])).
fof(f1142, plain, (spl48_94 <=> (e21 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl48_94])])).
fof(f3452, plain, (~ (e21 = op2(e22, e20)) | ~ spl48_110), inference(forward_demodulation, [], [f267, f1212])).
fof(f267, plain, ~ (op2(e21, e20) = op2(e22, e20)), inference(cnf_transformation, [], [f6])).
fof(f3451, plain, (~ spl48_96 | ~ spl48_308), inference(avatar_split_clause, [], [f3271, f2553, f1150])).
fof(f1150, plain, (spl48_96 <=> (e23 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl48_96])])).
fof(f3271, plain, (~ (e23 = op2(e22, e20)) | ~ spl48_308), inference(forward_demodulation, [], [f1848, f2554])).
fof(f1848, plain, ~ (op2(e22, e20) = h1(e10)), inference(backward_demodulation, [], [f265, f603])).
fof(f265, plain, ~ (op2(e20, e20) = op2(e22, e20)), inference(cnf_transformation, [], [f6])).
fof(f3448, plain, (~ spl48_91 | ~ spl48_123), inference(avatar_split_clause, [], [f3447, f1265, f1129])).
fof(f3447, plain, (~ (e22 = op2(e22, e21)) | ~ spl48_123), inference(forward_demodulation, [], [f271, f1267])).
fof(f271, plain, ~ (op2(e20, e21) = op2(e22, e21)), inference(cnf_transformation, [], [f6])).
fof(f3403, plain, (~ spl48_35 | ~ spl48_39), inference(avatar_split_clause, [], [f3402, f876, f859])).
fof(f859, plain, (spl48_35 <=> (e12 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl48_35])])).
fof(f3402, plain, (~ (e12 = op1(e11, e13)) | ~ spl48_39), inference(backward_demodulation, [], [f251, f878])).
fof(f251, plain, ~ (op1(e11, e12) = op1(e11, e13)), inference(cnf_transformation, [], [f5])).
fof(f3399, plain, (~ spl48_38 | ~ spl48_46), inference(avatar_split_clause, [], [f3393, f906, f872])).
fof(f872, plain, (spl48_38 <=> (e11 = op1(e11, e12))), introduced(avatar_definition, [new_symbols(naming, [spl48_38])])).
fof(f3393, plain, (~ (e11 = op1(e11, e12)) | ~ spl48_46), inference(backward_demodulation, [], [f247, f908])).
fof(f247, plain, ~ (op1(e11, e10) = op1(e11, e12)), inference(cnf_transformation, [], [f5])).
fof(f3398, plain, (~ spl48_30 | ~ spl48_46), inference(avatar_split_clause, [], [f3392, f906, f838])).
fof(f838, plain, (spl48_30 <=> (e11 = op1(e12, e10))), introduced(avatar_definition, [new_symbols(naming, [spl48_30])])).
fof(f3392, plain, (~ (e11 = op1(e12, e10)) | ~ spl48_46), inference(backward_demodulation, [], [f219, f908])).
fof(f219, plain, ~ (op1(e11, e10) = op1(e12, e10)), inference(cnf_transformation, [], [f5])).
fof(f3391, plain, (~ spl48_35 | ~ spl48_51), inference(avatar_split_clause, [], [f3387, f927, f859])).
fof(f3387, plain, (~ (e12 = op1(e11, e13)) | ~ spl48_51), inference(backward_demodulation, [], [f234, f929])).
fof(f234, plain, ~ (op1(e10, e13) = op1(e11, e13)), inference(cnf_transformation, [], [f5])).
fof(f3370, plain, (~ spl48_48 | ~ spl48_64), inference(avatar_split_clause, [], [f3365, f982, f914])).
fof(f914, plain, (spl48_48 <=> (e13 = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl48_48])])).
fof(f3365, plain, (~ (e13 = op1(e11, e10)) | ~ spl48_64), inference(backward_demodulation, [], [f216, f984])).
fof(f216, plain, ~ (op1(e10, e10) = op1(e11, e10)), inference(cnf_transformation, [], [f5])).
fof(f3364, plain, (spl48_232 | ~ spl48_65 | ~ spl48_308), inference(avatar_split_clause, [], [f3363, f2553, f1019, f2067])).
fof(f1019, plain, (spl48_65 <=> (e20 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl48_65])])).
fof(f3363, plain, ((e20 = h1(e13)) | (~ spl48_65 | ~ spl48_308)), inference(backward_demodulation, [], [f3164, f1021])).
fof(f1021, plain, ((e20 = op2(e23, e23)) | ~ spl48_65), inference(avatar_component_clause, [], [f1019])).
fof(f3357, plain, (~ spl48_72 | ~ spl48_76), inference(avatar_split_clause, [], [f3355, f1065, f1048])).
fof(f1048, plain, (spl48_72 <=> (e23 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl48_72])])).
fof(f3355, plain, (~ (e23 = op2(e23, e22)) | ~ spl48_76), inference(backward_demodulation, [], [f309, f1067])).
fof(f309, plain, ~ (op2(e23, e21) = op2(e23, e22)), inference(cnf_transformation, [], [f6])).
fof(f3335, plain, (~ spl48_117 | ~ spl48_119), inference(avatar_contradiction_clause, [], [f3334])).
fof(f3334, plain, ($false | (~ spl48_117 | ~ spl48_119)), inference(subsumption_resolution, [], [f3333, f319])).
fof(f3333, plain, ((e20 = e22) | (~ spl48_117 | ~ spl48_119)), inference(backward_demodulation, [], [f1250, f1242])).
fof(f1250, plain, ((e22 = op2(e20, e22)) | ~ spl48_119), inference(avatar_component_clause, [], [f1248])).
fof(f1248, plain, (spl48_119 <=> (e22 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl48_119])])).
fof(f3308, plain, (~ spl48_216 | ~ spl48_90), inference(avatar_split_clause, [], [f3207, f1125, f1981])).
fof(f1981, plain, (spl48_216 <=> (e21 = h3(e10))), introduced(avatar_definition, [new_symbols(naming, [spl48_216])])).
fof(f1125, plain, (spl48_90 <=> (e21 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl48_90])])).
fof(f3207, plain, (~ (e21 = h3(e10)) | ~ spl48_90), inference(backward_demodulation, [], [f1871, f1127])).
fof(f1127, plain, ((e21 = op2(e22, e21)) | ~ spl48_90), inference(avatar_component_clause, [], [f1125])).
fof(f1871, plain, ~ (op2(e22, e21) = h3(e10)), inference(backward_demodulation, [], [f303, f611])).
fof(f303, plain, ~ (op2(e22, e21) = op2(e22, e22)), inference(cnf_transformation, [], [f6])).
fof(f3307, plain, (spl48_216 | ~ spl48_86), inference(avatar_split_clause, [], [f3211, f1108, f1981])).
fof(f3211, plain, ((e21 = h3(e10)) | ~ spl48_86), inference(backward_demodulation, [], [f611, f1110])).
fof(f3304, plain, (~ spl48_120 | ~ spl48_308), inference(avatar_split_clause, [], [f3303, f2553, f1252])).
fof(f1252, plain, (spl48_120 <=> (e23 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl48_120])])).
fof(f3303, plain, (~ (e23 = op2(e20, e22)) | ~ spl48_308), inference(forward_demodulation, [], [f1851, f2554])).
fof(f1851, plain, ~ (op2(e20, e22) = h1(e10)), inference(backward_demodulation, [], [f289, f603])).
fof(f289, plain, ~ (op2(e20, e20) = op2(e20, e22)), inference(cnf_transformation, [], [f6])).
fof(f3291, plain, (~ spl48_116 | ~ spl48_308), inference(avatar_split_clause, [], [f3290, f2553, f1235])).
fof(f1235, plain, (spl48_116 <=> (e23 = op2(e20, e23))), introduced(avatar_definition, [new_symbols(naming, [spl48_116])])).
fof(f3290, plain, (~ (e23 = op2(e20, e23)) | ~ spl48_308), inference(forward_demodulation, [], [f1852, f2554])).
fof(f1852, plain, ~ (op2(e20, e23) = h1(e10)), inference(backward_demodulation, [], [f290, f603])).
fof(f290, plain, ~ (op2(e20, e20) = op2(e20, e23)), inference(cnf_transformation, [], [f6])).
fof(f3266, plain, (~ spl48_80 | ~ spl48_308), inference(avatar_split_clause, [], [f3265, f2553, f1082])).
fof(f1082, plain, (spl48_80 <=> (e23 = op2(e23, e20))), introduced(avatar_definition, [new_symbols(naming, [spl48_80])])).
fof(f3265, plain, (~ (e23 = op2(e23, e20)) | ~ spl48_308), inference(forward_demodulation, [], [f1849, f2554])).
fof(f1849, plain, ~ (op2(e23, e20) = h1(e10)), inference(backward_demodulation, [], [f266, f603])).
fof(f266, plain, ~ (op2(e20, e20) = op2(e23, e20)), inference(cnf_transformation, [], [f6])).
fof(f3262, plain, (spl48_79 | ~ spl48_222), inference(avatar_split_clause, [], [f2963, f2018, f1078])).
fof(f2963, plain, ((e22 = op2(e23, e20)) | ~ spl48_222), inference(backward_demodulation, [], [f1865, f2019])).
fof(f1865, plain, (e22 = op2(e23, h2(e10))), inference(backward_demodulation, [], [f1844, f607])).
fof(f607, plain, (op2(e21, e21) = h2(e10)), inference(cnf_transformation, [], [f15])).
fof(f3239, plain, (~ spl48_27 | ~ spl48_28), inference(avatar_contradiction_clause, [], [f3238])).
fof(f3238, plain, ($false | (~ spl48_27 | ~ spl48_28)), inference(subsumption_resolution, [], [f3237, f317])).
fof(f317, plain, ~ (e12 = e13), inference(cnf_transformation, [], [f7])).
fof(f3237, plain, ((e12 = e13) | (~ spl48_27 | ~ spl48_28)), inference(forward_demodulation, [], [f831, f827])).
fof(f3223, plain, (spl48_228 | ~ spl48_66 | ~ spl48_308), inference(avatar_split_clause, [], [f3222, f2553, f1023, f2047])).
fof(f1023, plain, (spl48_66 <=> (e21 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl48_66])])).
fof(f3222, plain, ((e21 = h1(e13)) | (~ spl48_66 | ~ spl48_308)), inference(backward_demodulation, [], [f3164, f1025])).
fof(f1025, plain, ((e21 = op2(e23, e23)) | ~ spl48_66), inference(avatar_component_clause, [], [f1023])).
fof(f3198, plain, (~ spl48_103 | ~ spl48_104), inference(avatar_contradiction_clause, [], [f3197])).
fof(f3197, plain, ($false | (~ spl48_103 | ~ spl48_104)), inference(subsumption_resolution, [], [f3196, f323])).
fof(f3196, plain, ((e22 = e23) | (~ spl48_103 | ~ spl48_104)), inference(forward_demodulation, [], [f1186, f1182])).
fof(f1186, plain, ((e23 = op2(e21, e22)) | ~ spl48_104), inference(avatar_component_clause, [], [f1184])).
fof(f1184, plain, (spl48_104 <=> (e23 = op2(e21, e22))), introduced(avatar_definition, [new_symbols(naming, [spl48_104])])).
fof(f3167, plain, (~ spl48_124 | ~ spl48_308), inference(avatar_split_clause, [], [f3163, f2553, f1269])).
fof(f1269, plain, (spl48_124 <=> (e23 = op2(e20, e21))), introduced(avatar_definition, [new_symbols(naming, [spl48_124])])).
fof(f3163, plain, (~ (e23 = op2(e20, e21)) | ~ spl48_308), inference(backward_demodulation, [], [f1850, f2554])).
fof(f1850, plain, ~ (op2(e20, e21) = h1(e10)), inference(backward_demodulation, [], [f288, f603])).
fof(f288, plain, ~ (op2(e20, e20) = op2(e20, e21)), inference(cnf_transformation, [], [f6])).
fof(f3143, plain, (~ spl48_97 | ~ spl48_222), inference(avatar_split_clause, [], [f3142, f2018, f1155])).
fof(f1155, plain, (spl48_97 <=> (e20 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl48_97])])).
fof(f3142, plain, (~ (e20 = op2(e21, e23)) | ~ spl48_222), inference(forward_demodulation, [], [f1863, f2019])).
fof(f1863, plain, ~ (op2(e21, e23) = h2(e10)), inference(backward_demodulation, [], [f298, f607])).
fof(f298, plain, ~ (op2(e21, e21) = op2(e21, e23)), inference(cnf_transformation, [], [f6])).
fof(f3139, plain, (~ spl48_98 | ~ spl48_110), inference(avatar_split_clause, [], [f3138, f1210, f1159])).
fof(f1159, plain, (spl48_98 <=> (e21 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl48_98])])).
fof(f3138, plain, (~ (e21 = op2(e21, e23)) | ~ spl48_110), inference(forward_demodulation, [], [f296, f1212])).
fof(f296, plain, ~ (op2(e21, e20) = op2(e21, e23)), inference(cnf_transformation, [], [f6])).
fof(f3137, plain, (~ spl48_91 | ~ spl48_27 | spl48_279), inference(avatar_split_clause, [], [f3136, f2368, f825, f1129])).
fof(f2368, plain, (spl48_279 <=> (op2(e22, e21) = h2(op1(e12, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl48_279])])).
fof(f3136, plain, (~ (e22 = op2(e22, e21)) | (~ spl48_27 | spl48_279)), inference(forward_demodulation, [], [f3135, f1857])).
fof(f3135, plain, (~ (op2(e22, e21) = h2(e12)) | (~ spl48_27 | spl48_279)), inference(forward_demodulation, [], [f2370, f827])).
fof(f2370, plain, (~ (op2(e22, e21) = h2(op1(e12, e11))) | spl48_279), inference(avatar_component_clause, [], [f2368])).
fof(f3134, plain, (~ spl48_100 | ~ spl48_84), inference(avatar_split_clause, [], [f3132, f1099, f1167])).
fof(f1099, plain, (spl48_84 <=> (e23 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl48_84])])).
fof(f3132, plain, (~ (e23 = op2(e21, e23)) | ~ spl48_84), inference(backward_demodulation, [], [f285, f1101])).
fof(f1101, plain, ((e23 = op2(e22, e23)) | ~ spl48_84), inference(avatar_component_clause, [], [f1099])).
fof(f285, plain, ~ (op2(e21, e23) = op2(e22, e23)), inference(cnf_transformation, [], [f6])).
fof(f3112, plain, (spl48_64 | ~ spl48_41), inference(avatar_split_clause, [], [f2758, f885, f982])).
fof(f2758, plain, ((op1(e10, e10) = e13) | ~ spl48_41), inference(backward_demodulation, [], [f598, f887])).
fof(f598, plain, (e13 = op1(op1(e11, e11), op1(e11, e11))), inference(cnf_transformation, [], [f12])).
fof(f12, plain, ((e13 = op1(op1(e11, e11), op1(e11, e11))) & (e12 = op1(op1(op1(e11, e11), op1(e11, e11)), op1(e11, e11))) & (e10 = op1(e11, e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG126+1.p', ax12)).
fof(f3094, plain, (~ spl48_20 | ~ spl48_24), inference(avatar_split_clause, [], [f3093, f812, f795])).
fof(f3093, plain, (~ (e13 = op1(e12, e13)) | ~ spl48_24), inference(forward_demodulation, [], [f257, f814])).
fof(f257, plain, ~ (op1(e12, e12) = op1(e12, e13)), inference(cnf_transformation, [], [f5])).
fof(f3070, plain, (~ spl48_13 | ~ spl48_15), inference(avatar_contradiction_clause, [], [f3069])).
fof(f3069, plain, ($false | (~ spl48_13 | ~ spl48_15)), inference(subsumption_resolution, [], [f3068, f313])).
fof(f3068, plain, ((e10 = e12) | (~ spl48_13 | ~ spl48_15)), inference(backward_demodulation, [], [f776, f768])).
fof(f768, plain, ((e10 = op1(e13, e10)) | ~ spl48_13), inference(avatar_component_clause, [], [f766])).
fof(f766, plain, (spl48_13 <=> (e10 = op1(e13, e10))), introduced(avatar_definition, [new_symbols(naming, [spl48_13])])).
fof(f3042, plain, (~ spl48_41 | ~ spl48_44), inference(avatar_contradiction_clause, [], [f3041])).
fof(f3041, plain, ($false | (~ spl48_41 | ~ spl48_44)), inference(subsumption_resolution, [], [f3040, f314])).
fof(f314, plain, ~ (e10 = e13), inference(cnf_transformation, [], [f7])).
fof(f3040, plain, ((e10 = e13) | (~ spl48_41 | ~ spl48_44)), inference(forward_demodulation, [], [f899, f887])).
fof(f899, plain, ((e13 = op1(e11, e11)) | ~ spl48_44), inference(avatar_component_clause, [], [f897])).
fof(f897, plain, (spl48_44 <=> (e13 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl48_44])])).
fof(f3032, plain, (~ spl48_53 | ~ spl48_55), inference(avatar_contradiction_clause, [], [f3031])).
fof(f3031, plain, ($false | (~ spl48_53 | ~ spl48_55)), inference(subsumption_resolution, [], [f3030, f313])).
fof(f3030, plain, ((e10 = e12) | (~ spl48_53 | ~ spl48_55)), inference(forward_demodulation, [], [f946, f938])).
fof(f946, plain, ((e12 = op1(e10, e12)) | ~ spl48_55), inference(avatar_component_clause, [], [f944])).
fof(f944, plain, (spl48_55 <=> (e12 = op1(e10, e12))), introduced(avatar_definition, [new_symbols(naming, [spl48_55])])).
fof(f3028, plain, (~ spl48_26 | ~ spl48_58), inference(avatar_split_clause, [], [f3027, f957, f821])).
fof(f821, plain, (spl48_26 <=> (e11 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl48_26])])).
fof(f3027, plain, (~ (e11 = op1(e12, e11)) | ~ spl48_58), inference(backward_demodulation, [], [f223, f959])).
fof(f223, plain, ~ (op1(e10, e11) = op1(e12, e11)), inference(cnf_transformation, [], [f5])).
fof(f3015, plain, (~ spl48_88 | spl48_272), inference(avatar_contradiction_clause, [], [f3014])).
fof(f3014, plain, ($false | (~ spl48_88 | spl48_272)), inference(subsumption_resolution, [], [f3013, f2312])).
fof(f2312, plain, (~ (e23 = h3(e10)) | spl48_272), inference(avatar_component_clause, [], [f2310])).
fof(f3013, plain, ((e23 = h3(e10)) | ~ spl48_88), inference(backward_demodulation, [], [f611, f1118])).
fof(f1118, plain, ((e23 = op2(e22, e22)) | ~ spl48_88), inference(avatar_component_clause, [], [f1116])).
fof(f1116, plain, (spl48_88 <=> (e23 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl48_88])])).
fof(f3012, plain, (~ spl48_212 | ~ spl48_91), inference(avatar_split_clause, [], [f3010, f1129, f1961])).
fof(f3010, plain, (~ (e22 = h3(e10)) | ~ spl48_91), inference(backward_demodulation, [], [f1871, f1131])).
fof(f3003, plain, (~ spl48_105 | ~ spl48_108), inference(avatar_contradiction_clause, [], [f3002])).
fof(f3002, plain, ($false | (~ spl48_105 | ~ spl48_108)), inference(subsumption_resolution, [], [f3001, f320])).
fof(f320, plain, ~ (e20 = e23), inference(cnf_transformation, [], [f8])).
fof(f3001, plain, ((e20 = e23) | (~ spl48_105 | ~ spl48_108)), inference(forward_demodulation, [], [f1203, f1191])).
fof(f1203, plain, ((e23 = op2(e21, e21)) | ~ spl48_108), inference(avatar_component_clause, [], [f1201])).
fof(f1201, plain, (spl48_108 <=> (e23 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl48_108])])).
fof(f2996, plain, (~ spl48_308 | ~ spl48_112), inference(avatar_split_clause, [], [f2994, f1218, f2553])).
fof(f1218, plain, (spl48_112 <=> (e23 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl48_112])])).
fof(f2994, plain, (~ (e23 = h1(e10)) | ~ spl48_112), inference(backward_demodulation, [], [f1847, f1220])).
fof(f1220, plain, ((e23 = op2(e21, e20)) | ~ spl48_112), inference(avatar_component_clause, [], [f1218])).
fof(f1847, plain, ~ (op2(e21, e20) = h1(e10)), inference(backward_demodulation, [], [f264, f603])).
fof(f264, plain, ~ (op2(e20, e20) = op2(e21, e20)), inference(cnf_transformation, [], [f6])).
fof(f2979, plain, (~ spl48_90 | ~ spl48_122), inference(avatar_split_clause, [], [f2976, f1261, f1125])).
fof(f2976, plain, (~ (e21 = op2(e22, e21)) | ~ spl48_122), inference(backward_demodulation, [], [f271, f1263])).
fof(f2974, plain, (~ spl48_125 | spl48_234), inference(avatar_contradiction_clause, [], [f2973])).
fof(f2973, plain, ($false | (~ spl48_125 | spl48_234)), inference(subsumption_resolution, [], [f2972, f2081])).
fof(f2081, plain, (~ (e20 = h1(e10)) | spl48_234), inference(avatar_component_clause, [], [f2079])).
fof(f2079, plain, (spl48_234 <=> (e20 = h1(e10))), introduced(avatar_definition, [new_symbols(naming, [spl48_234])])).
fof(f2972, plain, ((e20 = h1(e10)) | ~ spl48_125), inference(backward_demodulation, [], [f603, f1276])).
fof(f1276, plain, ((e20 = op2(e20, e20)) | ~ spl48_125), inference(avatar_component_clause, [], [f1274])).
fof(f1274, plain, (spl48_125 <=> (e20 = op2(e20, e20))), introduced(avatar_definition, [new_symbols(naming, [spl48_125])])).
fof(f2969, plain, (spl48_128 | ~ spl48_222), inference(avatar_split_clause, [], [f2962, f2018, f1286])).
fof(f1286, plain, (spl48_128 <=> (op2(e20, e20) = e23)), introduced(avatar_definition, [new_symbols(naming, [spl48_128])])).
fof(f2962, plain, ((op2(e20, e20) = e23) | ~ spl48_222), inference(backward_demodulation, [], [f1864, f2019])).
fof(f1864, plain, (e23 = op2(h2(e10), h2(e10))), inference(backward_demodulation, [], [f601, f607])).
fof(f2967, plain, (~ spl48_109 | ~ spl48_222), inference(avatar_split_clause, [], [f2960, f2018, f1206])).
fof(f1206, plain, (spl48_109 <=> (e20 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl48_109])])).
fof(f2960, plain, (~ (e20 = op2(e21, e20)) | ~ spl48_222), inference(backward_demodulation, [], [f1861, f2019])).
fof(f1861, plain, ~ (op2(e21, e20) = h2(e10)), inference(backward_demodulation, [], [f294, f607])).
fof(f294, plain, ~ (op2(e21, e20) = op2(e21, e21)), inference(cnf_transformation, [], [f6])).
fof(f2965, plain, (~ spl48_121 | ~ spl48_222), inference(avatar_split_clause, [], [f2958, f2018, f1257])).
fof(f1257, plain, (spl48_121 <=> (e20 = op2(e20, e21))), introduced(avatar_definition, [new_symbols(naming, [spl48_121])])).
fof(f2958, plain, (~ (e20 = op2(e20, e21)) | ~ spl48_222), inference(backward_demodulation, [], [f1858, f2019])).
fof(f1858, plain, ~ (op2(e20, e21) = h2(e10)), inference(backward_demodulation, [], [f270, f607])).
fof(f270, plain, ~ (op2(e20, e21) = op2(e21, e21)), inference(cnf_transformation, [], [f6])).
fof(f2953, plain, (~ spl48_122 | ~ spl48_74), inference(avatar_split_clause, [], [f2952, f1057, f1261])).
fof(f1057, plain, (spl48_74 <=> (e21 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl48_74])])).
fof(f2952, plain, (~ (e21 = op2(e20, e21)) | ~ spl48_74), inference(forward_demodulation, [], [f272, f1059])).
fof(f1059, plain, ((e21 = op2(e23, e21)) | ~ spl48_74), inference(avatar_component_clause, [], [f1057])).
fof(f272, plain, ~ (op2(e20, e21) = op2(e23, e21)), inference(cnf_transformation, [], [f6])).
fof(f2943, plain, (~ spl48_111 | ~ spl48_79), inference(avatar_split_clause, [], [f2942, f1078, f1214])).
fof(f1214, plain, (spl48_111 <=> (e22 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl48_111])])).
fof(f2942, plain, (~ (e22 = op2(e21, e20)) | ~ spl48_79), inference(forward_demodulation, [], [f268, f1080])).
fof(f268, plain, ~ (op2(e21, e20) = op2(e23, e20)), inference(cnf_transformation, [], [f6])).
fof(f2913, plain, (~ spl48_57 | ~ spl48_41), inference(avatar_split_clause, [], [f2912, f885, f953])).
fof(f953, plain, (spl48_57 <=> (e10 = op1(e10, e11))), introduced(avatar_definition, [new_symbols(naming, [spl48_57])])).
fof(f2912, plain, (~ (e10 = op1(e10, e11)) | ~ spl48_41), inference(forward_demodulation, [], [f222, f887])).
fof(f222, plain, ~ (op1(e10, e11) = op1(e11, e11)), inference(cnf_transformation, [], [f5])).
fof(f2899, plain, (~ spl48_47 | ~ spl48_15), inference(avatar_split_clause, [], [f2898, f774, f910])).
fof(f910, plain, (spl48_47 <=> (e12 = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl48_47])])).
fof(f2898, plain, (~ (e12 = op1(e11, e10)) | ~ spl48_15), inference(forward_demodulation, [], [f220, f776])).
fof(f220, plain, ~ (op1(e11, e10) = op1(e13, e10)), inference(cnf_transformation, [], [f5])).
fof(f2897, plain, (~ spl48_45 | ~ spl48_41), inference(avatar_split_clause, [], [f2896, f885, f902])).
fof(f902, plain, (spl48_45 <=> (e10 = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl48_45])])).
fof(f2896, plain, (~ (e10 = op1(e11, e10)) | ~ spl48_41), inference(forward_demodulation, [], [f246, f887])).
fof(f246, plain, ~ (op1(e11, e10) = op1(e11, e11)), inference(cnf_transformation, [], [f5])).
fof(f2884, plain, (~ spl48_29 | ~ spl48_17), inference(avatar_split_clause, [], [f2883, f783, f834])).
fof(f783, plain, (spl48_17 <=> (e10 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl48_17])])).
fof(f2883, plain, (~ (e10 = op1(e12, e10)) | ~ spl48_17), inference(forward_demodulation, [], [f254, f785])).
fof(f785, plain, ((e10 = op1(e12, e13)) | ~ spl48_17), inference(avatar_component_clause, [], [f783])).
fof(f254, plain, ~ (op1(e12, e10) = op1(e12, e13)), inference(cnf_transformation, [], [f5])).
fof(f2833, plain, (~ spl48_3 | ~ spl48_15), inference(avatar_split_clause, [], [f2827, f774, f723])).
fof(f723, plain, (spl48_3 <=> (e12 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl48_3])])).
fof(f2827, plain, (~ (e12 = op1(e13, e13)) | ~ spl48_15), inference(backward_demodulation, [], [f260, f776])).
fof(f260, plain, ~ (op1(e13, e10) = op1(e13, e13)), inference(cnf_transformation, [], [f5])).
fof(f2832, plain, (~ spl48_7 | ~ spl48_15), inference(avatar_split_clause, [], [f2826, f774, f740])).
fof(f740, plain, (spl48_7 <=> (e12 = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl48_7])])).
fof(f2826, plain, (~ (e12 = op1(e13, e12)) | ~ spl48_15), inference(backward_demodulation, [], [f259, f776])).
fof(f259, plain, ~ (op1(e13, e10) = op1(e13, e12)), inference(cnf_transformation, [], [f5])).
fof(f2831, plain, (~ spl48_11 | ~ spl48_15), inference(avatar_split_clause, [], [f2825, f774, f757])).
fof(f757, plain, (spl48_11 <=> (e12 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl48_11])])).
fof(f2825, plain, (~ (e12 = op1(e13, e11)) | ~ spl48_15), inference(backward_demodulation, [], [f258, f776])).
fof(f258, plain, ~ (op1(e13, e10) = op1(e13, e11)), inference(cnf_transformation, [], [f5])).
fof(f2765, plain, (spl48_15 | ~ spl48_41), inference(avatar_split_clause, [], [f2760, f885, f774])).
fof(f2760, plain, ((e12 = op1(e13, e10)) | ~ spl48_41), inference(backward_demodulation, [], [f1842, f887])).
fof(f1842, plain, (e12 = op1(e13, op1(e11, e11))), inference(forward_demodulation, [], [f597, f598])).
fof(f597, plain, (e12 = op1(op1(op1(e11, e11), op1(e11, e11)), op1(e11, e11))), inference(cnf_transformation, [], [f12])).
fof(f2764, plain, (~ spl48_33 | ~ spl48_41), inference(avatar_split_clause, [], [f2757, f885, f851])).
fof(f851, plain, (spl48_33 <=> (e10 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl48_33])])).
fof(f2757, plain, (~ (e10 = op1(e11, e13)) | ~ spl48_41), inference(backward_demodulation, [], [f250, f887])).
fof(f250, plain, ~ (op1(e11, e11) = op1(e11, e13)), inference(cnf_transformation, [], [f5])).
fof(f2762, plain, (~ spl48_9 | ~ spl48_41), inference(avatar_split_clause, [], [f2755, f885, f749])).
fof(f749, plain, (spl48_9 <=> (e10 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl48_9])])).
fof(f2755, plain, (~ (e10 = op1(e13, e11)) | ~ spl48_41), inference(backward_demodulation, [], [f226, f887])).
fof(f226, plain, ~ (op1(e11, e11) = op1(e13, e11)), inference(cnf_transformation, [], [f5])).
fof(f2761, plain, (~ spl48_25 | ~ spl48_41), inference(avatar_split_clause, [], [f2754, f885, f817])).
fof(f817, plain, (spl48_25 <=> (e10 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl48_25])])).
fof(f2754, plain, (~ (e10 = op1(e12, e11)) | ~ spl48_41), inference(backward_demodulation, [], [f225, f887])).
fof(f225, plain, ~ (op1(e11, e11) = op1(e12, e11)), inference(cnf_transformation, [], [f5])).
fof(f2734, plain, (~ spl48_49 | ~ spl48_53), inference(avatar_split_clause, [], [f2729, f936, f919])).
fof(f919, plain, (spl48_49 <=> (e10 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl48_49])])).
fof(f2729, plain, (~ (e10 = op1(e10, e13)) | ~ spl48_53), inference(backward_demodulation, [], [f245, f938])).
fof(f245, plain, ~ (op1(e10, e12) = op1(e10, e13)), inference(cnf_transformation, [], [f5])).
fof(f2733, plain, (~ spl48_5 | ~ spl48_53), inference(avatar_split_clause, [], [f2728, f936, f732])).
fof(f732, plain, (spl48_5 <=> (e10 = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl48_5])])).
fof(f2728, plain, (~ (e10 = op1(e13, e12)) | ~ spl48_53), inference(backward_demodulation, [], [f230, f938])).
fof(f230, plain, ~ (op1(e10, e12) = op1(e13, e12)), inference(cnf_transformation, [], [f5])).
fof(f2732, plain, (~ spl48_21 | ~ spl48_53), inference(avatar_split_clause, [], [f2727, f936, f800])).
fof(f800, plain, (spl48_21 <=> (e10 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl48_21])])).
fof(f2727, plain, (~ (e10 = op1(e12, e12)) | ~ spl48_53), inference(backward_demodulation, [], [f229, f938])).
fof(f229, plain, ~ (op1(e10, e12) = op1(e12, e12)), inference(cnf_transformation, [], [f5])).
fof(f2715, plain, (~ spl48_52 | ~ spl48_64), inference(avatar_split_clause, [], [f2709, f982, f931])).
fof(f931, plain, (spl48_52 <=> (e13 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl48_52])])).
fof(f2709, plain, (~ (e13 = op1(e10, e13)) | ~ spl48_64), inference(backward_demodulation, [], [f242, f984])).
fof(f242, plain, ~ (op1(e10, e10) = op1(e10, e13)), inference(cnf_transformation, [], [f5])).
fof(f2713, plain, (~ spl48_60 | ~ spl48_64), inference(avatar_split_clause, [], [f2707, f982, f965])).
fof(f965, plain, (spl48_60 <=> (e13 = op1(e10, e11))), introduced(avatar_definition, [new_symbols(naming, [spl48_60])])).
fof(f2707, plain, (~ (e13 = op1(e10, e11)) | ~ spl48_64), inference(backward_demodulation, [], [f240, f984])).
fof(f240, plain, ~ (op1(e10, e10) = op1(e10, e11)), inference(cnf_transformation, [], [f5])).
fof(f2712, plain, (~ spl48_16 | ~ spl48_64), inference(avatar_split_clause, [], [f2706, f982, f778])).
fof(f778, plain, (spl48_16 <=> (e13 = op1(e13, e10))), introduced(avatar_definition, [new_symbols(naming, [spl48_16])])).
fof(f2706, plain, (~ (e13 = op1(e13, e10)) | ~ spl48_64), inference(backward_demodulation, [], [f218, f984])).
fof(f218, plain, ~ (op1(e10, e10) = op1(e13, e10)), inference(cnf_transformation, [], [f5])).
fof(f2678, plain, (~ spl48_78 | ~ spl48_79), inference(avatar_contradiction_clause, [], [f2677])).
fof(f2677, plain, ($false | (~ spl48_78 | ~ spl48_79)), inference(subsumption_resolution, [], [f2676, f321])).
fof(f2676, plain, ((e21 = e22) | (~ spl48_78 | ~ spl48_79)), inference(backward_demodulation, [], [f1080, f1076])).
fof(f1076, plain, ((e21 = op2(e23, e20)) | ~ spl48_78), inference(avatar_component_clause, [], [f1074])).
fof(f1074, plain, (spl48_78 <=> (e21 = op2(e23, e20))), introduced(avatar_definition, [new_symbols(naming, [spl48_78])])).
fof(f2674, plain, (~ spl48_75 | ~ spl48_79), inference(avatar_split_clause, [], [f2670, f1078, f1061])).
fof(f1061, plain, (spl48_75 <=> (e22 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl48_75])])).
fof(f2670, plain, (~ (e22 = op2(e23, e21)) | ~ spl48_79), inference(backward_demodulation, [], [f306, f1080])).
fof(f306, plain, ~ (op2(e23, e20) = op2(e23, e21)), inference(cnf_transformation, [], [f6])).
fof(f2664, plain, (~ spl48_83 | ~ spl48_84), inference(avatar_contradiction_clause, [], [f2663])).
fof(f2663, plain, ($false | (~ spl48_83 | ~ spl48_84)), inference(subsumption_resolution, [], [f2662, f323])).
fof(f2662, plain, ((e22 = e23) | (~ spl48_83 | ~ spl48_84)), inference(backward_demodulation, [], [f1101, f1097])).
fof(f1097, plain, ((e22 = op2(e22, e23)) | ~ spl48_83), inference(avatar_component_clause, [], [f1095])).
fof(f2656, plain, (spl48_220 | ~ spl48_85), inference(avatar_split_clause, [], [f2655, f1104, f2001])).
fof(f2001, plain, (spl48_220 <=> (e20 = h3(e10))), introduced(avatar_definition, [new_symbols(naming, [spl48_220])])).
fof(f1104, plain, (spl48_85 <=> (e20 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl48_85])])).
fof(f2655, plain, ((e20 = h3(e10)) | ~ spl48_85), inference(backward_demodulation, [], [f611, f1106])).
fof(f1106, plain, ((e20 = op2(e22, e22)) | ~ spl48_85), inference(avatar_component_clause, [], [f1104])).
fof(f2639, plain, (~ spl48_220 | ~ spl48_93), inference(avatar_split_clause, [], [f2634, f1138, f2001])).
fof(f2634, plain, (~ (e20 = h3(e10)) | ~ spl48_93), inference(backward_demodulation, [], [f1870, f1140])).
fof(f1870, plain, ~ (op2(e22, e20) = h3(e10)), inference(backward_demodulation, [], [f301, f611])).
fof(f301, plain, ~ (op2(e22, e20) = op2(e22, e22)), inference(cnf_transformation, [], [f6])).
fof(f2638, plain, (~ spl48_234 | ~ spl48_93), inference(avatar_split_clause, [], [f2633, f1138, f2079])).
fof(f2633, plain, (~ (e20 = h1(e10)) | ~ spl48_93), inference(backward_demodulation, [], [f1848, f1140])).
fof(f2607, plain, (spl48_222 | ~ spl48_105), inference(avatar_split_clause, [], [f2606, f1189, f2018])).
fof(f2606, plain, ((e20 = h2(e10)) | ~ spl48_105), inference(backward_demodulation, [], [f607, f1191])).
fof(f2588, plain, (~ spl48_99 | ~ spl48_115), inference(avatar_split_clause, [], [f2584, f1231, f1163])).
fof(f1163, plain, (spl48_99 <=> (e22 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl48_99])])).
fof(f2584, plain, (~ (e22 = op2(e21, e23)) | ~ spl48_115), inference(backward_demodulation, [], [f282, f1233])).
fof(f282, plain, ~ (op2(e20, e23) = op2(e21, e23)), inference(cnf_transformation, [], [f6])).
fof(f2581, plain, (~ spl48_113 | ~ spl48_117), inference(avatar_split_clause, [], [f2576, f1240, f1223])).
fof(f1223, plain, (spl48_113 <=> (e20 = op2(e20, e23))), introduced(avatar_definition, [new_symbols(naming, [spl48_113])])).
fof(f2576, plain, (~ (e20 = op2(e20, e23)) | ~ spl48_117), inference(backward_demodulation, [], [f293, f1242])).
fof(f293, plain, ~ (op2(e20, e22) = op2(e20, e23)), inference(cnf_transformation, [], [f6])).
fof(f2580, plain, (~ spl48_69 | ~ spl48_117), inference(avatar_split_clause, [], [f2575, f1240, f1036])).
fof(f1036, plain, (spl48_69 <=> (e20 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl48_69])])).
fof(f2575, plain, (~ (e20 = op2(e23, e22)) | ~ spl48_117), inference(backward_demodulation, [], [f278, f1242])).
fof(f278, plain, ~ (op2(e20, e22) = op2(e23, e22)), inference(cnf_transformation, [], [f6])).
fof(f2559, plain, (spl48_308 | ~ spl48_128), inference(avatar_split_clause, [], [f2558, f1286, f2553])).
fof(f2558, plain, ((e23 = h1(e10)) | ~ spl48_128), inference(backward_demodulation, [], [f603, f1288])).
fof(f1288, plain, ((op2(e20, e20) = e23) | ~ spl48_128), inference(avatar_component_clause, [], [f1286])).
fof(f2407, plain, (spl48_221 | ~ spl48_273 | ~ spl48_274 | ~ spl48_275 | ~ spl48_276 | ~ spl48_277 | ~ spl48_278 | ~ spl48_279 | ~ spl48_280 | ~ spl48_281 | ~ spl48_282 | ~ spl48_283 | ~ spl48_284 | ~ spl48_285 | ~ spl48_286 | ~ spl48_287 | ~ spl48_288), inference(avatar_split_clause, [], [f2342, f2404, f2400, f2396, f2392, f2388, f2384, f2380, f2376, f2372, f2368, f2364, f2360, f2356, f2352, f2348, f2344, f2014])).
fof(f2014, plain, (spl48_221 <=> sP39), introduced(avatar_definition, [new_symbols(naming, [spl48_221])])).
fof(f2342, plain, (~ (e23 = h2(op1(e10, e10))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), e21)) | ~ (h2(op1(e10, e12)) = op2(h2(e10), e22)) | ~ (h2(op1(e10, e13)) = op2(h2(e10), e23)) | ~ (h2(op1(e11, e10)) = op2(e21, h2(e10))) | ~ (h2(e10) = h2(op1(e11, e11))) | ~ (op2(e21, e22) = h2(op1(e11, e12))) | ~ (op2(e21, e23) = h2(op1(e11, e13))) | ~ (h2(op1(e12, e10)) = op2(e22, h2(e10))) | ~ (op2(e22, e21) = h2(op1(e12, e11))) | ~ (h3(e10) = h2(op1(e12, e12))) | ~ (op2(e22, e23) = h2(op1(e12, e13))) | ~ (e22 = h2(op1(e13, e10))) | ~ (op2(e23, e21) = h2(op1(e13, e11))) | ~ (op2(e23, e22) = h2(op1(e13, e12))) | ~ (h4(e10) = h2(op1(e13, e13))) | sP39), inference(forward_demodulation, [], [f2341, f1864])).
fof(f2341, plain, (~ (h2(op1(e10, e11)) = op2(h2(e10), e21)) | ~ (h2(op1(e10, e12)) = op2(h2(e10), e22)) | ~ (h2(op1(e10, e13)) = op2(h2(e10), e23)) | ~ (h2(op1(e11, e10)) = op2(e21, h2(e10))) | ~ (h2(e10) = h2(op1(e11, e11))) | ~ (op2(e21, e22) = h2(op1(e11, e12))) | ~ (op2(e21, e23) = h2(op1(e11, e13))) | ~ (h2(op1(e12, e10)) = op2(e22, h2(e10))) | ~ (op2(e22, e21) = h2(op1(e12, e11))) | ~ (h3(e10) = h2(op1(e12, e12))) | ~ (op2(e22, e23) = h2(op1(e12, e13))) | ~ (e22 = h2(op1(e13, e10))) | ~ (op2(e23, e21) = h2(op1(e13, e11))) | ~ (op2(e23, e22) = h2(op1(e13, e12))) | ~ (h4(e10) = h2(op1(e13, e13))) | sP39 | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f2340, f606])).
fof(f2340, plain, (~ (h2(op1(e10, e12)) = op2(h2(e10), e22)) | ~ (h2(op1(e10, e13)) = op2(h2(e10), e23)) | ~ (h2(op1(e11, e10)) = op2(e21, h2(e10))) | ~ (h2(e10) = h2(op1(e11, e11))) | ~ (op2(e21, e22) = h2(op1(e11, e12))) | ~ (op2(e21, e23) = h2(op1(e11, e13))) | ~ (h2(op1(e12, e10)) = op2(e22, h2(e10))) | ~ (op2(e22, e21) = h2(op1(e12, e11))) | ~ (h3(e10) = h2(op1(e12, e12))) | ~ (op2(e22, e23) = h2(op1(e12, e13))) | ~ (e22 = h2(op1(e13, e10))) | ~ (op2(e23, e21) = h2(op1(e13, e11))) | ~ (op2(e23, e22) = h2(op1(e13, e12))) | ~ (h4(e10) = h2(op1(e13, e13))) | sP39 | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f2339, f1857])).
fof(f2339, plain, (~ (h2(op1(e10, e13)) = op2(h2(e10), e23)) | ~ (h2(op1(e11, e10)) = op2(e21, h2(e10))) | ~ (h2(e10) = h2(op1(e11, e11))) | ~ (op2(e21, e22) = h2(op1(e11, e12))) | ~ (op2(e21, e23) = h2(op1(e11, e13))) | ~ (h2(op1(e12, e10)) = op2(e22, h2(e10))) | ~ (op2(e22, e21) = h2(op1(e12, e11))) | ~ (h3(e10) = h2(op1(e12, e12))) | ~ (op2(e22, e23) = h2(op1(e12, e13))) | ~ (e22 = h2(op1(e13, e10))) | ~ (op2(e23, e21) = h2(op1(e13, e11))) | ~ (op2(e23, e22) = h2(op1(e13, e12))) | ~ (h4(e10) = h2(op1(e13, e13))) | sP39 | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f2338, f1855])).
fof(f2338, plain, (~ (h2(op1(e11, e10)) = op2(e21, h2(e10))) | ~ (h2(e10) = h2(op1(e11, e11))) | ~ (op2(e21, e22) = h2(op1(e11, e12))) | ~ (op2(e21, e23) = h2(op1(e11, e13))) | ~ (h2(op1(e12, e10)) = op2(e22, h2(e10))) | ~ (op2(e22, e21) = h2(op1(e12, e11))) | ~ (h3(e10) = h2(op1(e12, e12))) | ~ (op2(e22, e23) = h2(op1(e12, e13))) | ~ (e22 = h2(op1(e13, e10))) | ~ (op2(e23, e21) = h2(op1(e13, e11))) | ~ (op2(e23, e22) = h2(op1(e13, e12))) | ~ (h4(e10) = h2(op1(e13, e13))) | sP39 | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f2337, f606])).
fof(f2337, plain, (~ (h2(e10) = h2(op1(e11, e11))) | ~ (op2(e21, e22) = h2(op1(e11, e12))) | ~ (op2(e21, e23) = h2(op1(e11, e13))) | ~ (h2(op1(e12, e10)) = op2(e22, h2(e10))) | ~ (op2(e22, e21) = h2(op1(e12, e11))) | ~ (h3(e10) = h2(op1(e12, e12))) | ~ (op2(e22, e23) = h2(op1(e12, e13))) | ~ (e22 = h2(op1(e13, e10))) | ~ (op2(e23, e21) = h2(op1(e13, e11))) | ~ (op2(e23, e22) = h2(op1(e13, e12))) | ~ (h4(e10) = h2(op1(e13, e13))) | sP39 | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f2336, f607])).
fof(f2336, plain, (~ (op2(e21, e21) = h2(op1(e11, e11))) | ~ (op2(e21, e22) = h2(op1(e11, e12))) | ~ (op2(e21, e23) = h2(op1(e11, e13))) | ~ (h2(op1(e12, e10)) = op2(e22, h2(e10))) | ~ (op2(e22, e21) = h2(op1(e12, e11))) | ~ (h3(e10) = h2(op1(e12, e12))) | ~ (op2(e22, e23) = h2(op1(e12, e13))) | ~ (e22 = h2(op1(e13, e10))) | ~ (op2(e23, e21) = h2(op1(e13, e11))) | ~ (op2(e23, e22) = h2(op1(e13, e12))) | ~ (h4(e10) = h2(op1(e13, e13))) | sP39 | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f2335, f606])).
fof(f2335, plain, (~ (op2(e21, e22) = h2(op1(e11, e12))) | ~ (op2(e21, e23) = h2(op1(e11, e13))) | ~ (h2(op1(e12, e10)) = op2(e22, h2(e10))) | ~ (op2(e22, e21) = h2(op1(e12, e11))) | ~ (h3(e10) = h2(op1(e12, e12))) | ~ (op2(e22, e23) = h2(op1(e12, e13))) | ~ (e22 = h2(op1(e13, e10))) | ~ (op2(e23, e21) = h2(op1(e13, e11))) | ~ (op2(e23, e22) = h2(op1(e13, e12))) | ~ (h4(e10) = h2(op1(e13, e13))) | sP39 | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f2334, f606])).
fof(f2334, plain, (~ (h2(op1(e11, e12)) = op2(h2(e11), e22)) | ~ (op2(e21, e23) = h2(op1(e11, e13))) | ~ (h2(op1(e12, e10)) = op2(e22, h2(e10))) | ~ (op2(e22, e21) = h2(op1(e12, e11))) | ~ (h3(e10) = h2(op1(e12, e12))) | ~ (op2(e22, e23) = h2(op1(e12, e13))) | ~ (e22 = h2(op1(e13, e10))) | ~ (op2(e23, e21) = h2(op1(e13, e11))) | ~ (op2(e23, e22) = h2(op1(e13, e12))) | ~ (h4(e10) = h2(op1(e13, e13))) | sP39 | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f2333, f1857])).
fof(f2333, plain, (~ (op2(e21, e23) = h2(op1(e11, e13))) | ~ (h2(op1(e12, e10)) = op2(e22, h2(e10))) | ~ (op2(e22, e21) = h2(op1(e12, e11))) | ~ (h3(e10) = h2(op1(e12, e12))) | ~ (op2(e22, e23) = h2(op1(e12, e13))) | ~ (e22 = h2(op1(e13, e10))) | ~ (op2(e23, e21) = h2(op1(e13, e11))) | ~ (op2(e23, e22) = h2(op1(e13, e12))) | ~ (h4(e10) = h2(op1(e13, e13))) | sP39 | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f2332, f606])).
fof(f2332, plain, (~ (h2(op1(e11, e13)) = op2(h2(e11), e23)) | ~ (h2(op1(e12, e10)) = op2(e22, h2(e10))) | ~ (op2(e22, e21) = h2(op1(e12, e11))) | ~ (h3(e10) = h2(op1(e12, e12))) | ~ (op2(e22, e23) = h2(op1(e12, e13))) | ~ (e22 = h2(op1(e13, e10))) | ~ (op2(e23, e21) = h2(op1(e13, e11))) | ~ (op2(e23, e22) = h2(op1(e13, e12))) | ~ (h4(e10) = h2(op1(e13, e13))) | sP39 | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f2331, f1855])).
fof(f2331, plain, (~ (h2(op1(e12, e10)) = op2(e22, h2(e10))) | ~ (op2(e22, e21) = h2(op1(e12, e11))) | ~ (h3(e10) = h2(op1(e12, e12))) | ~ (op2(e22, e23) = h2(op1(e12, e13))) | ~ (e22 = h2(op1(e13, e10))) | ~ (op2(e23, e21) = h2(op1(e13, e11))) | ~ (op2(e23, e22) = h2(op1(e13, e12))) | ~ (h4(e10) = h2(op1(e13, e13))) | sP39 | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f2330, f1857])).
fof(f2330, plain, (~ (op2(e22, e21) = h2(op1(e12, e11))) | ~ (h3(e10) = h2(op1(e12, e12))) | ~ (op2(e22, e23) = h2(op1(e12, e13))) | ~ (e22 = h2(op1(e13, e10))) | ~ (op2(e23, e21) = h2(op1(e13, e11))) | ~ (op2(e23, e22) = h2(op1(e13, e12))) | ~ (h4(e10) = h2(op1(e13, e13))) | sP39 | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f2329, f1857])).
fof(f2329, plain, (~ (h2(op1(e12, e11)) = op2(h2(e12), e21)) | ~ (h3(e10) = h2(op1(e12, e12))) | ~ (op2(e22, e23) = h2(op1(e12, e13))) | ~ (e22 = h2(op1(e13, e10))) | ~ (op2(e23, e21) = h2(op1(e13, e11))) | ~ (op2(e23, e22) = h2(op1(e13, e12))) | ~ (h4(e10) = h2(op1(e13, e13))) | sP39 | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f2328, f606])).
fof(f2328, plain, (~ (h3(e10) = h2(op1(e12, e12))) | ~ (op2(e22, e23) = h2(op1(e12, e13))) | ~ (e22 = h2(op1(e13, e10))) | ~ (op2(e23, e21) = h2(op1(e13, e11))) | ~ (op2(e23, e22) = h2(op1(e13, e12))) | ~ (h4(e10) = h2(op1(e13, e13))) | sP39 | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f2327, f611])).
fof(f2327, plain, (~ (op2(e22, e22) = h2(op1(e12, e12))) | ~ (op2(e22, e23) = h2(op1(e12, e13))) | ~ (e22 = h2(op1(e13, e10))) | ~ (op2(e23, e21) = h2(op1(e13, e11))) | ~ (op2(e23, e22) = h2(op1(e13, e12))) | ~ (h4(e10) = h2(op1(e13, e13))) | sP39 | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f2326, f1857])).
fof(f2326, plain, (~ (op2(e22, e23) = h2(op1(e12, e13))) | ~ (e22 = h2(op1(e13, e10))) | ~ (op2(e23, e21) = h2(op1(e13, e11))) | ~ (op2(e23, e22) = h2(op1(e13, e12))) | ~ (h4(e10) = h2(op1(e13, e13))) | sP39 | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f2325, f1857])).
fof(f2325, plain, (~ (h2(op1(e12, e13)) = op2(h2(e12), e23)) | ~ (e22 = h2(op1(e13, e10))) | ~ (op2(e23, e21) = h2(op1(e13, e11))) | ~ (op2(e23, e22) = h2(op1(e13, e12))) | ~ (h4(e10) = h2(op1(e13, e13))) | sP39 | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f2324, f1855])).
fof(f2324, plain, (~ (e22 = h2(op1(e13, e10))) | ~ (op2(e23, e21) = h2(op1(e13, e11))) | ~ (op2(e23, e22) = h2(op1(e13, e12))) | ~ (h4(e10) = h2(op1(e13, e13))) | sP39 | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f2323, f1865])).
fof(f2323, plain, (~ (h2(op1(e13, e10)) = op2(e23, h2(e10))) | ~ (op2(e23, e21) = h2(op1(e13, e11))) | ~ (op2(e23, e22) = h2(op1(e13, e12))) | ~ (h4(e10) = h2(op1(e13, e13))) | sP39 | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f2322, f1855])).
fof(f2322, plain, (~ (op2(e23, e21) = h2(op1(e13, e11))) | ~ (op2(e23, e22) = h2(op1(e13, e12))) | ~ (h4(e10) = h2(op1(e13, e13))) | sP39 | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f2321, f1855])).
fof(f2321, plain, (~ (h2(op1(e13, e11)) = op2(h2(e13), e21)) | ~ (op2(e23, e22) = h2(op1(e13, e12))) | ~ (h4(e10) = h2(op1(e13, e13))) | sP39 | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f2320, f606])).
fof(f2320, plain, (~ (op2(e23, e22) = h2(op1(e13, e12))) | ~ (h4(e10) = h2(op1(e13, e13))) | sP39 | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f2319, f1855])).
fof(f2319, plain, (~ (h2(op1(e13, e12)) = op2(h2(e13), e22)) | ~ (h4(e10) = h2(op1(e13, e13))) | sP39 | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f2318, f1857])).
fof(f2318, plain, (~ (h4(e10) = h2(op1(e13, e13))) | sP39 | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f2317, f615])).
fof(f2317, plain, (~ (op2(e23, e23) = h2(op1(e13, e13))) | sP39 | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f2316, f1855])).
fof(f2316, plain, (sP39 | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(subsumption_resolution, [], [f2315, f2009])).
fof(f2009, plain, ~ sP40, inference(subsumption_resolution, [], [f647, f606])).
fof(f647, plain, (~ (e21 = h2(e11)) | ~ sP40), inference(cnf_transformation, [], [f115])).
fof(f115, plain, ((~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10))) | ~ sP40), inference(nnf_transformation, [], [f63])).
fof(f63, plain, ((~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10))) | ~ sP40), inference(usedef, [], [e63])).
fof(e63, plain, (sP40 <=> (~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP40])])).
fof(f2315, plain, (sP40 | sP39 | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(subsumption_resolution, [], [f2314, f2006])).
fof(f2006, plain, ~ sP41, inference(subsumption_resolution, [], [f644, f1857])).
fof(f644, plain, (~ (e22 = h2(e12)) | ~ sP41), inference(cnf_transformation, [], [f114])).
fof(f114, plain, ((~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10))) | ~ sP41), inference(nnf_transformation, [], [f64])).
fof(f64, plain, ((~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10))) | ~ sP41), inference(usedef, [], [e64])).
fof(e64, plain, (sP41 <=> (~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP41])])).
fof(f2314, plain, (sP41 | sP40 | sP39 | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(subsumption_resolution, [], [f673, f1855])).
fof(f673, plain, (~ (e23 = h2(e13)) | sP41 | sP40 | sP39 | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(cnf_transformation, [], [f71])).
fof(f71, plain, (((~ (e23 = h4(e13)) & ~ (e23 = h4(e12)) & ~ (e23 = h4(e11)) & ~ (e23 = h4(e10))) | sP47 | sP46 | sP45 | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) & ((~ (e23 = h3(e13)) & ~ (e23 = h3(e12)) & ~ (e23 = h3(e11)) & ~ (e23 = h3(e10))) | sP44 | sP43 | sP42 | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) & ((~ (e23 = h2(e13)) & ~ (e23 = h2(e12)) & ~ (e23 = h2(e11)) & ~ (e23 = h2(e10))) | sP41 | sP40 | sP39 | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) & ((~ (e23 = h1(e13)) & ~ (e23 = h1(e12)) & ~ (e23 = h1(e11)) & ~ (e23 = h1(e10))) | sP38 | sP37 | sP36 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(definition_folding, [], [f20, e70, e69, e68, e67, e66, e65, e64, e63, e62, e61, e60, e59])).
fof(f59, plain, ((~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10))) | ~ sP36), inference(usedef, [], [e59])).
fof(e59, plain, (sP36 <=> (~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP36])])).
fof(f60, plain, ((~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10))) | ~ sP37), inference(usedef, [], [e60])).
fof(e60, plain, (sP37 <=> (~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP37])])).
fof(f61, plain, ((~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10))) | ~ sP38), inference(usedef, [], [e61])).
fof(e61, plain, (sP38 <=> (~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP38])])).
fof(f62, plain, ((~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10))) | ~ sP39), inference(usedef, [], [e62])).
fof(e62, plain, (sP39 <=> (~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP39])])).
fof(f65, plain, ((~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10))) | ~ sP42), inference(usedef, [], [e65])).
fof(e65, plain, (sP42 <=> (~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP42])])).
fof(f66, plain, ((~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10))) | ~ sP43), inference(usedef, [], [e66])).
fof(e66, plain, (sP43 <=> (~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP43])])).
fof(f67, plain, ((~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10))) | ~ sP44), inference(usedef, [], [e67])).
fof(e67, plain, (sP44 <=> (~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP44])])).
fof(f68, plain, ((~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10))) | ~ sP45), inference(usedef, [], [e68])).
fof(e68, plain, (sP45 <=> (~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP45])])).
fof(f69, plain, ((~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10))) | ~ sP46), inference(usedef, [], [e69])).
fof(e69, plain, (sP46 <=> (~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP46])])).
fof(f70, plain, ((~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10))) | ~ sP47), inference(usedef, [], [e70])).
fof(e70, plain, (sP47 <=> (~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP47])])).
fof(f20, plain, (((~ (e23 = h4(e13)) & ~ (e23 = h4(e12)) & ~ (e23 = h4(e11)) & ~ (e23 = h4(e10))) | (~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10))) | (~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10))) | (~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10))) | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) & ((~ (e23 = h3(e13)) & ~ (e23 = h3(e12)) & ~ (e23 = h3(e11)) & ~ (e23 = h3(e10))) | (~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10))) | (~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10))) | (~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10))) | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) & ((~ (e23 = h2(e13)) & ~ (e23 = h2(e12)) & ~ (e23 = h2(e11)) & ~ (e23 = h2(e10))) | (~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10))) | (~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10))) | (~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) & ((~ (e23 = h1(e13)) & ~ (e23 = h1(e12)) & ~ (e23 = h1(e11)) & ~ (e23 = h1(e10))) | (~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10))) | (~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10))) | (~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10))) | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(ennf_transformation, [], [f19])).
fof(f19, plain, ~ ((((e23 = h4(e13)) | (e23 = h4(e12)) | (e23 = h4(e11)) | (e23 = h4(e10))) & ((e22 = h4(e13)) | (e22 = h4(e12)) | (e22 = h4(e11)) | (e22 = h4(e10))) & ((e21 = h4(e13)) | (e21 = h4(e12)) | (e21 = h4(e11)) | (e21 = h4(e10))) & ((e20 = h4(e13)) | (e20 = h4(e12)) | (e20 = h4(e11)) | (e20 = h4(e10))) & (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) & (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) & (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) & (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) & (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) & (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) & (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) & (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) & (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) & (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) & (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) & (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) & (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) & (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) & (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) & (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) | (((e23 = h3(e13)) | (e23 = h3(e12)) | (e23 = h3(e11)) | (e23 = h3(e10))) & ((e22 = h3(e13)) | (e22 = h3(e12)) | (e22 = h3(e11)) | (e22 = h3(e10))) & ((e21 = h3(e13)) | (e21 = h3(e12)) | (e21 = h3(e11)) | (e21 = h3(e10))) & ((e20 = h3(e13)) | (e20 = h3(e12)) | (e20 = h3(e11)) | (e20 = h3(e10))) & (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) & (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) & (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) & (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) & (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) & (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) & (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) & (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) & (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) & (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) & (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) & (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) & (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) & (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) & (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) & (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) | (((e23 = h2(e13)) | (e23 = h2(e12)) | (e23 = h2(e11)) | (e23 = h2(e10))) & ((e22 = h2(e13)) | (e22 = h2(e12)) | (e22 = h2(e11)) | (e22 = h2(e10))) & ((e21 = h2(e13)) | (e21 = h2(e12)) | (e21 = h2(e11)) | (e21 = h2(e10))) & ((e20 = h2(e13)) | (e20 = h2(e12)) | (e20 = h2(e11)) | (e20 = h2(e10))) & (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) & (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) & (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) & (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) & (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) & (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) & (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) & (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) & (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) & (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) & (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) & (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) & (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) & (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) & (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) & (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) | (((e23 = h1(e13)) | (e23 = h1(e12)) | (e23 = h1(e11)) | (e23 = h1(e10))) & ((e22 = h1(e13)) | (e22 = h1(e12)) | (e22 = h1(e11)) | (e22 = h1(e10))) & ((e21 = h1(e13)) | (e21 = h1(e12)) | (e21 = h1(e11)) | (e21 = h1(e10))) & ((e20 = h1(e13)) | (e20 = h1(e12)) | (e20 = h1(e11)) | (e20 = h1(e10))) & (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) & (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) & (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) & (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) & (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) & (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) & (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) & (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) & (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) & (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) & (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) & (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) & (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) & (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) & (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) & (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(negated_conjecture, [], [f18])).
fof(f18, plain, ~ ((((e23 = h4(e13)) | (e23 = h4(e12)) | (e23 = h4(e11)) | (e23 = h4(e10))) & ((e22 = h4(e13)) | (e22 = h4(e12)) | (e22 = h4(e11)) | (e22 = h4(e10))) & ((e21 = h4(e13)) | (e21 = h4(e12)) | (e21 = h4(e11)) | (e21 = h4(e10))) & ((e20 = h4(e13)) | (e20 = h4(e12)) | (e20 = h4(e11)) | (e20 = h4(e10))) & (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) & (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) & (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) & (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) & (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) & (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) & (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) & (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) & (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) & (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) & (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) & (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) & (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) & (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) & (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) & (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) | (((e23 = h3(e13)) | (e23 = h3(e12)) | (e23 = h3(e11)) | (e23 = h3(e10))) & ((e22 = h3(e13)) | (e22 = h3(e12)) | (e22 = h3(e11)) | (e22 = h3(e10))) & ((e21 = h3(e13)) | (e21 = h3(e12)) | (e21 = h3(e11)) | (e21 = h3(e10))) & ((e20 = h3(e13)) | (e20 = h3(e12)) | (e20 = h3(e11)) | (e20 = h3(e10))) & (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) & (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) & (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) & (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) & (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) & (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) & (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) & (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) & (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) & (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) & (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) & (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) & (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) & (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) & (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) & (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) | (((e23 = h2(e13)) | (e23 = h2(e12)) | (e23 = h2(e11)) | (e23 = h2(e10))) & ((e22 = h2(e13)) | (e22 = h2(e12)) | (e22 = h2(e11)) | (e22 = h2(e10))) & ((e21 = h2(e13)) | (e21 = h2(e12)) | (e21 = h2(e11)) | (e21 = h2(e10))) & ((e20 = h2(e13)) | (e20 = h2(e12)) | (e20 = h2(e11)) | (e20 = h2(e10))) & (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) & (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) & (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) & (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) & (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) & (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) & (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) & (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) & (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) & (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) & (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) & (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) & (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) & (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) & (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) & (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) | (((e23 = h1(e13)) | (e23 = h1(e12)) | (e23 = h1(e11)) | (e23 = h1(e10))) & ((e22 = h1(e13)) | (e22 = h1(e12)) | (e22 = h1(e11)) | (e22 = h1(e10))) & ((e21 = h1(e13)) | (e21 = h1(e12)) | (e21 = h1(e11)) | (e21 = h1(e10))) & ((e20 = h1(e13)) | (e20 = h1(e12)) | (e20 = h1(e11)) | (e20 = h1(e10))) & (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) & (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) & (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) & (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) & (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) & (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) & (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) & (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) & (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) & (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) & (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) & (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) & (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) & (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) & (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) & (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG126+1.p', co1)).
fof(f2021, plain, (~ spl48_221 | ~ spl48_222), inference(avatar_split_clause, [], [f650, f2018, f2014])).
fof(f650, plain, (~ (e20 = h2(e10)) | ~ sP39), inference(cnf_transformation, [], [f116])).
fof(f116, plain, ((~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10))) | ~ sP39), inference(nnf_transformation, [], [f62])).
fof(f1845, plain, spl48_105, inference(avatar_split_clause, [], [f599, f1189])).
fof(f599, plain, (e20 = op2(e21, e21)), inference(cnf_transformation, [], [f13])).
fof(f1843, plain, spl48_41, inference(avatar_split_clause, [], [f596, f885])).
fof(f596, plain, (e10 = op1(e11, e11)), inference(cnf_transformation, [], [f12])).
fof(f1821, plain, (spl48_188 | spl48_183 | spl48_178 | ~ spl48_128 | spl48_113), inference(avatar_split_clause, [], [f586, f1223, f1286, f1733, f1760, f1787])).
fof(f1787, plain, (spl48_188 <=> sP18), introduced(avatar_definition, [new_symbols(naming, [spl48_188])])).
fof(f1760, plain, (spl48_183 <=> sP19), introduced(avatar_definition, [new_symbols(naming, [spl48_183])])).
fof(f1733, plain, (spl48_178 <=> sP20), introduced(avatar_definition, [new_symbols(naming, [spl48_178])])).
fof(f586, plain, ((e20 = op2(e20, e23)) | ~ (op2(e20, e20) = e23) | sP20 | sP19 | sP18), inference(cnf_transformation, [], [f58])).
fof(f58, plain, ((((~ (e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23))) & (~ (e22 = op2(e23, e23)) | ~ (e23 = op2(e23, e22))) & (~ (e21 = op2(e23, e23)) | ~ (e23 = op2(e23, e21))) & (~ (e20 = op2(e23, e23)) | ~ (e23 = op2(e23, e20))) & (e23 = op2(e23, e23)) & (e23 = op2(e23, e23))) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21) & ((((e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23))) & ((e22 = op2(e22, e23)) | ~ (e23 = op2(e22, e22))) & ((e21 = op2(e21, e23)) | ~ (e23 = op2(e21, e21))) & ((e20 = op2(e20, e23)) | ~ (op2(e20, e20) = e23)) & (e23 = op2(e23, op2(e23, e23))) & (e22 = op2(e23, op2(e23, e22))) & (e21 = op2(e23, op2(e23, e21))) & (e20 = op2(e23, op2(e23, e20)))) | sP20 | sP19 | sP18)), inference(definition_folding, [], [f11, e57, e56, e55, e54, e53, e52, e51, e50, e49, e48, e47, e46, e45, e44, e43, e42, e41, e40])).
fof(f40, plain, ((((e23 = op2(e23, e20)) | ~ (e20 = op2(e23, e23))) & ((e22 = op2(e22, e20)) | ~ (e20 = op2(e22, e22))) & ((e21 = op2(e21, e20)) | ~ (e20 = op2(e21, e21))) & ((e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20))) & (e23 = op2(e20, op2(e20, e23))) & (e22 = op2(e20, op2(e20, e22))) & (e21 = op2(e20, op2(e20, e21))) & (e20 = op2(e20, op2(e20, e20)))) | ~ sP18), inference(usedef, [], [e40])).
fof(e40, plain, (sP18 <=> (((e23 = op2(e23, e20)) | ~ (e20 = op2(e23, e23))) & ((e22 = op2(e22, e20)) | ~ (e20 = op2(e22, e22))) & ((e21 = op2(e21, e20)) | ~ (e20 = op2(e21, e21))) & ((e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20))) & (e23 = op2(e20, op2(e20, e23))) & (e22 = op2(e20, op2(e20, e22))) & (e21 = op2(e20, op2(e20, e21))) & (e20 = op2(e20, op2(e20, e20))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f41, plain, ((((e23 = op2(e23, e21)) | ~ (e21 = op2(e23, e23))) & ((e22 = op2(e22, e21)) | ~ (e21 = op2(e22, e22))) & ((e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & ((e20 = op2(e20, e21)) | ~ (op2(e20, e20) = e21)) & (e23 = op2(e21, op2(e21, e23))) & (e22 = op2(e21, op2(e21, e22))) & (e21 = op2(e21, op2(e21, e21))) & (e20 = op2(e21, op2(e21, e20)))) | ~ sP19), inference(usedef, [], [e41])).
fof(e41, plain, (sP19 <=> (((e23 = op2(e23, e21)) | ~ (e21 = op2(e23, e23))) & ((e22 = op2(e22, e21)) | ~ (e21 = op2(e22, e22))) & ((e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & ((e20 = op2(e20, e21)) | ~ (op2(e20, e20) = e21)) & (e23 = op2(e21, op2(e21, e23))) & (e22 = op2(e21, op2(e21, e22))) & (e21 = op2(e21, op2(e21, e21))) & (e20 = op2(e21, op2(e21, e20))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP19])])).
fof(f42, plain, ((((e23 = op2(e23, e22)) | ~ (e22 = op2(e23, e23))) & ((e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & ((e21 = op2(e21, e22)) | ~ (e22 = op2(e21, e21))) & ((e20 = op2(e20, e22)) | ~ (op2(e20, e20) = e22)) & (e23 = op2(e22, op2(e22, e23))) & (e22 = op2(e22, op2(e22, e22))) & (e21 = op2(e22, op2(e22, e21))) & (e20 = op2(e22, op2(e22, e20)))) | ~ sP20), inference(usedef, [], [e42])).
fof(e42, plain, (sP20 <=> (((e23 = op2(e23, e22)) | ~ (e22 = op2(e23, e23))) & ((e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & ((e21 = op2(e21, e22)) | ~ (e22 = op2(e21, e21))) & ((e20 = op2(e20, e22)) | ~ (op2(e20, e20) = e22)) & (e23 = op2(e22, op2(e22, e23))) & (e22 = op2(e22, op2(e22, e22))) & (e21 = op2(e22, op2(e22, e21))) & (e20 = op2(e22, op2(e22, e20))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP20])])).
fof(f43, plain, (((~ (op2(e20, e20) = e23) | ~ (e20 = op2(e20, e23))) & (~ (op2(e20, e20) = e22) | ~ (e20 = op2(e20, e22))) & (~ (op2(e20, e20) = e21) | ~ (e20 = op2(e20, e21))) & (~ (e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20))) & (e20 = op2(e20, e20)) & (e20 = op2(e20, e20))) | ~ sP21), inference(usedef, [], [e43])).
fof(e43, plain, (sP21 <=> ((~ (op2(e20, e20) = e23) | ~ (e20 = op2(e20, e23))) & (~ (op2(e20, e20) = e22) | ~ (e20 = op2(e20, e22))) & (~ (op2(e20, e20) = e21) | ~ (e20 = op2(e20, e21))) & (~ (e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20))) & (e20 = op2(e20, e20)) & (e20 = op2(e20, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP21])])).
fof(f44, plain, (((~ (e23 = op2(e21, e21)) | ~ (e21 = op2(e21, e23))) & (~ (e22 = op2(e21, e21)) | ~ (e21 = op2(e21, e22))) & (~ (e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & (~ (e20 = op2(e21, e21)) | ~ (e21 = op2(e21, e20))) & (e20 = op2(e21, e20)) & (e20 = op2(e20, e21))) | ~ sP22), inference(usedef, [], [e44])).
fof(e44, plain, (sP22 <=> ((~ (e23 = op2(e21, e21)) | ~ (e21 = op2(e21, e23))) & (~ (e22 = op2(e21, e21)) | ~ (e21 = op2(e21, e22))) & (~ (e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & (~ (e20 = op2(e21, e21)) | ~ (e21 = op2(e21, e20))) & (e20 = op2(e21, e20)) & (e20 = op2(e20, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP22])])).
fof(f45, plain, (((~ (e23 = op2(e22, e22)) | ~ (e22 = op2(e22, e23))) & (~ (e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & (~ (e21 = op2(e22, e22)) | ~ (e22 = op2(e22, e21))) & (~ (e20 = op2(e22, e22)) | ~ (e22 = op2(e22, e20))) & (e20 = op2(e22, e20)) & (e20 = op2(e20, e22))) | ~ sP23), inference(usedef, [], [e45])).
fof(e45, plain, (sP23 <=> ((~ (e23 = op2(e22, e22)) | ~ (e22 = op2(e22, e23))) & (~ (e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & (~ (e21 = op2(e22, e22)) | ~ (e22 = op2(e22, e21))) & (~ (e20 = op2(e22, e22)) | ~ (e22 = op2(e22, e20))) & (e20 = op2(e22, e20)) & (e20 = op2(e20, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP23])])).
fof(f46, plain, (((~ (e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23))) & (~ (e22 = op2(e23, e23)) | ~ (e23 = op2(e23, e22))) & (~ (e21 = op2(e23, e23)) | ~ (e23 = op2(e23, e21))) & (~ (e20 = op2(e23, e23)) | ~ (e23 = op2(e23, e20))) & (e20 = op2(e23, e20)) & (e20 = op2(e20, e23))) | ~ sP24), inference(usedef, [], [e46])).
fof(e46, plain, (sP24 <=> ((~ (e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23))) & (~ (e22 = op2(e23, e23)) | ~ (e23 = op2(e23, e22))) & (~ (e21 = op2(e23, e23)) | ~ (e23 = op2(e23, e21))) & (~ (e20 = op2(e23, e23)) | ~ (e23 = op2(e23, e20))) & (e20 = op2(e23, e20)) & (e20 = op2(e20, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP24])])).
fof(f47, plain, (((~ (op2(e20, e20) = e23) | ~ (e20 = op2(e20, e23))) & (~ (op2(e20, e20) = e22) | ~ (e20 = op2(e20, e22))) & (~ (op2(e20, e20) = e21) | ~ (e20 = op2(e20, e21))) & (~ (e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20))) & (e21 = op2(e20, e21)) & (e21 = op2(e21, e20))) | ~ sP25), inference(usedef, [], [e47])).
fof(e47, plain, (sP25 <=> ((~ (op2(e20, e20) = e23) | ~ (e20 = op2(e20, e23))) & (~ (op2(e20, e20) = e22) | ~ (e20 = op2(e20, e22))) & (~ (op2(e20, e20) = e21) | ~ (e20 = op2(e20, e21))) & (~ (e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20))) & (e21 = op2(e20, e21)) & (e21 = op2(e21, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP25])])).
fof(f48, plain, (((~ (e23 = op2(e21, e21)) | ~ (e21 = op2(e21, e23))) & (~ (e22 = op2(e21, e21)) | ~ (e21 = op2(e21, e22))) & (~ (e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & (~ (e20 = op2(e21, e21)) | ~ (e21 = op2(e21, e20))) & (e21 = op2(e21, e21)) & (e21 = op2(e21, e21))) | ~ sP26), inference(usedef, [], [e48])).
fof(e48, plain, (sP26 <=> ((~ (e23 = op2(e21, e21)) | ~ (e21 = op2(e21, e23))) & (~ (e22 = op2(e21, e21)) | ~ (e21 = op2(e21, e22))) & (~ (e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & (~ (e20 = op2(e21, e21)) | ~ (e21 = op2(e21, e20))) & (e21 = op2(e21, e21)) & (e21 = op2(e21, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP26])])).
fof(f49, plain, (((~ (e23 = op2(e22, e22)) | ~ (e22 = op2(e22, e23))) & (~ (e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & (~ (e21 = op2(e22, e22)) | ~ (e22 = op2(e22, e21))) & (~ (e20 = op2(e22, e22)) | ~ (e22 = op2(e22, e20))) & (e21 = op2(e22, e21)) & (e21 = op2(e21, e22))) | ~ sP27), inference(usedef, [], [e49])).
fof(e49, plain, (sP27 <=> ((~ (e23 = op2(e22, e22)) | ~ (e22 = op2(e22, e23))) & (~ (e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & (~ (e21 = op2(e22, e22)) | ~ (e22 = op2(e22, e21))) & (~ (e20 = op2(e22, e22)) | ~ (e22 = op2(e22, e20))) & (e21 = op2(e22, e21)) & (e21 = op2(e21, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP27])])).
fof(f50, plain, (((~ (e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23))) & (~ (e22 = op2(e23, e23)) | ~ (e23 = op2(e23, e22))) & (~ (e21 = op2(e23, e23)) | ~ (e23 = op2(e23, e21))) & (~ (e20 = op2(e23, e23)) | ~ (e23 = op2(e23, e20))) & (e21 = op2(e23, e21)) & (e21 = op2(e21, e23))) | ~ sP28), inference(usedef, [], [e50])).
fof(e50, plain, (sP28 <=> ((~ (e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23))) & (~ (e22 = op2(e23, e23)) | ~ (e23 = op2(e23, e22))) & (~ (e21 = op2(e23, e23)) | ~ (e23 = op2(e23, e21))) & (~ (e20 = op2(e23, e23)) | ~ (e23 = op2(e23, e20))) & (e21 = op2(e23, e21)) & (e21 = op2(e21, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP28])])).
fof(f51, plain, (((~ (op2(e20, e20) = e23) | ~ (e20 = op2(e20, e23))) & (~ (op2(e20, e20) = e22) | ~ (e20 = op2(e20, e22))) & (~ (op2(e20, e20) = e21) | ~ (e20 = op2(e20, e21))) & (~ (e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20))) & (e22 = op2(e20, e22)) & (e22 = op2(e22, e20))) | ~ sP29), inference(usedef, [], [e51])).
fof(e51, plain, (sP29 <=> ((~ (op2(e20, e20) = e23) | ~ (e20 = op2(e20, e23))) & (~ (op2(e20, e20) = e22) | ~ (e20 = op2(e20, e22))) & (~ (op2(e20, e20) = e21) | ~ (e20 = op2(e20, e21))) & (~ (e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20))) & (e22 = op2(e20, e22)) & (e22 = op2(e22, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP29])])).
fof(f52, plain, (((~ (e23 = op2(e21, e21)) | ~ (e21 = op2(e21, e23))) & (~ (e22 = op2(e21, e21)) | ~ (e21 = op2(e21, e22))) & (~ (e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & (~ (e20 = op2(e21, e21)) | ~ (e21 = op2(e21, e20))) & (e22 = op2(e21, e22)) & (e22 = op2(e22, e21))) | ~ sP30), inference(usedef, [], [e52])).
fof(e52, plain, (sP30 <=> ((~ (e23 = op2(e21, e21)) | ~ (e21 = op2(e21, e23))) & (~ (e22 = op2(e21, e21)) | ~ (e21 = op2(e21, e22))) & (~ (e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & (~ (e20 = op2(e21, e21)) | ~ (e21 = op2(e21, e20))) & (e22 = op2(e21, e22)) & (e22 = op2(e22, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP30])])).
fof(f53, plain, (((~ (e23 = op2(e22, e22)) | ~ (e22 = op2(e22, e23))) & (~ (e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & (~ (e21 = op2(e22, e22)) | ~ (e22 = op2(e22, e21))) & (~ (e20 = op2(e22, e22)) | ~ (e22 = op2(e22, e20))) & (e22 = op2(e22, e22)) & (e22 = op2(e22, e22))) | ~ sP31), inference(usedef, [], [e53])).
fof(e53, plain, (sP31 <=> ((~ (e23 = op2(e22, e22)) | ~ (e22 = op2(e22, e23))) & (~ (e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & (~ (e21 = op2(e22, e22)) | ~ (e22 = op2(e22, e21))) & (~ (e20 = op2(e22, e22)) | ~ (e22 = op2(e22, e20))) & (e22 = op2(e22, e22)) & (e22 = op2(e22, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP31])])).
fof(f54, plain, (((~ (e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23))) & (~ (e22 = op2(e23, e23)) | ~ (e23 = op2(e23, e22))) & (~ (e21 = op2(e23, e23)) | ~ (e23 = op2(e23, e21))) & (~ (e20 = op2(e23, e23)) | ~ (e23 = op2(e23, e20))) & (e22 = op2(e23, e22)) & (e22 = op2(e22, e23))) | ~ sP32), inference(usedef, [], [e54])).
fof(e54, plain, (sP32 <=> ((~ (e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23))) & (~ (e22 = op2(e23, e23)) | ~ (e23 = op2(e23, e22))) & (~ (e21 = op2(e23, e23)) | ~ (e23 = op2(e23, e21))) & (~ (e20 = op2(e23, e23)) | ~ (e23 = op2(e23, e20))) & (e22 = op2(e23, e22)) & (e22 = op2(e22, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP32])])).
fof(f55, plain, (((~ (op2(e20, e20) = e23) | ~ (e20 = op2(e20, e23))) & (~ (op2(e20, e20) = e22) | ~ (e20 = op2(e20, e22))) & (~ (op2(e20, e20) = e21) | ~ (e20 = op2(e20, e21))) & (~ (e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20))) & (e23 = op2(e20, e23)) & (e23 = op2(e23, e20))) | ~ sP33), inference(usedef, [], [e55])).
fof(e55, plain, (sP33 <=> ((~ (op2(e20, e20) = e23) | ~ (e20 = op2(e20, e23))) & (~ (op2(e20, e20) = e22) | ~ (e20 = op2(e20, e22))) & (~ (op2(e20, e20) = e21) | ~ (e20 = op2(e20, e21))) & (~ (e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20))) & (e23 = op2(e20, e23)) & (e23 = op2(e23, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP33])])).
fof(f56, plain, (((~ (e23 = op2(e21, e21)) | ~ (e21 = op2(e21, e23))) & (~ (e22 = op2(e21, e21)) | ~ (e21 = op2(e21, e22))) & (~ (e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & (~ (e20 = op2(e21, e21)) | ~ (e21 = op2(e21, e20))) & (e23 = op2(e21, e23)) & (e23 = op2(e23, e21))) | ~ sP34), inference(usedef, [], [e56])).
fof(e56, plain, (sP34 <=> ((~ (e23 = op2(e21, e21)) | ~ (e21 = op2(e21, e23))) & (~ (e22 = op2(e21, e21)) | ~ (e21 = op2(e21, e22))) & (~ (e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & (~ (e20 = op2(e21, e21)) | ~ (e21 = op2(e21, e20))) & (e23 = op2(e21, e23)) & (e23 = op2(e23, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP34])])).
fof(f57, plain, (((~ (e23 = op2(e22, e22)) | ~ (e22 = op2(e22, e23))) & (~ (e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & (~ (e21 = op2(e22, e22)) | ~ (e22 = op2(e22, e21))) & (~ (e20 = op2(e22, e22)) | ~ (e22 = op2(e22, e20))) & (e23 = op2(e22, e23)) & (e23 = op2(e23, e22))) | ~ sP35), inference(usedef, [], [e57])).
fof(e57, plain, (sP35 <=> ((~ (e23 = op2(e22, e22)) | ~ (e22 = op2(e22, e23))) & (~ (e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & (~ (e21 = op2(e22, e22)) | ~ (e22 = op2(e22, e21))) & (~ (e20 = op2(e22, e22)) | ~ (e22 = op2(e22, e20))) & (e23 = op2(e22, e23)) & (e23 = op2(e23, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP35])])).
fof(f11, plain, ((((~ (e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23))) & (~ (e22 = op2(e23, e23)) | ~ (e23 = op2(e23, e22))) & (~ (e21 = op2(e23, e23)) | ~ (e23 = op2(e23, e21))) & (~ (e20 = op2(e23, e23)) | ~ (e23 = op2(e23, e20))) & (e23 = op2(e23, e23)) & (e23 = op2(e23, e23))) | ((~ (e23 = op2(e22, e22)) | ~ (e22 = op2(e22, e23))) & (~ (e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & (~ (e21 = op2(e22, e22)) | ~ (e22 = op2(e22, e21))) & (~ (e20 = op2(e22, e22)) | ~ (e22 = op2(e22, e20))) & (e23 = op2(e22, e23)) & (e23 = op2(e23, e22))) | ((~ (e23 = op2(e21, e21)) | ~ (e21 = op2(e21, e23))) & (~ (e22 = op2(e21, e21)) | ~ (e21 = op2(e21, e22))) & (~ (e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & (~ (e20 = op2(e21, e21)) | ~ (e21 = op2(e21, e20))) & (e23 = op2(e21, e23)) & (e23 = op2(e23, e21))) | ((~ (op2(e20, e20) = e23) | ~ (e20 = op2(e20, e23))) & (~ (op2(e20, e20) = e22) | ~ (e20 = op2(e20, e22))) & (~ (op2(e20, e20) = e21) | ~ (e20 = op2(e20, e21))) & (~ (e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20))) & (e23 = op2(e20, e23)) & (e23 = op2(e23, e20))) | ((~ (e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23))) & (~ (e22 = op2(e23, e23)) | ~ (e23 = op2(e23, e22))) & (~ (e21 = op2(e23, e23)) | ~ (e23 = op2(e23, e21))) & (~ (e20 = op2(e23, e23)) | ~ (e23 = op2(e23, e20))) & (e22 = op2(e23, e22)) & (e22 = op2(e22, e23))) | ((~ (e23 = op2(e22, e22)) | ~ (e22 = op2(e22, e23))) & (~ (e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & (~ (e21 = op2(e22, e22)) | ~ (e22 = op2(e22, e21))) & (~ (e20 = op2(e22, e22)) | ~ (e22 = op2(e22, e20))) & (e22 = op2(e22, e22)) & (e22 = op2(e22, e22))) | ((~ (e23 = op2(e21, e21)) | ~ (e21 = op2(e21, e23))) & (~ (e22 = op2(e21, e21)) | ~ (e21 = op2(e21, e22))) & (~ (e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & (~ (e20 = op2(e21, e21)) | ~ (e21 = op2(e21, e20))) & (e22 = op2(e21, e22)) & (e22 = op2(e22, e21))) | ((~ (op2(e20, e20) = e23) | ~ (e20 = op2(e20, e23))) & (~ (op2(e20, e20) = e22) | ~ (e20 = op2(e20, e22))) & (~ (op2(e20, e20) = e21) | ~ (e20 = op2(e20, e21))) & (~ (e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20))) & (e22 = op2(e20, e22)) & (e22 = op2(e22, e20))) | ((~ (e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23))) & (~ (e22 = op2(e23, e23)) | ~ (e23 = op2(e23, e22))) & (~ (e21 = op2(e23, e23)) | ~ (e23 = op2(e23, e21))) & (~ (e20 = op2(e23, e23)) | ~ (e23 = op2(e23, e20))) & (e21 = op2(e23, e21)) & (e21 = op2(e21, e23))) | ((~ (e23 = op2(e22, e22)) | ~ (e22 = op2(e22, e23))) & (~ (e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & (~ (e21 = op2(e22, e22)) | ~ (e22 = op2(e22, e21))) & (~ (e20 = op2(e22, e22)) | ~ (e22 = op2(e22, e20))) & (e21 = op2(e22, e21)) & (e21 = op2(e21, e22))) | ((~ (e23 = op2(e21, e21)) | ~ (e21 = op2(e21, e23))) & (~ (e22 = op2(e21, e21)) | ~ (e21 = op2(e21, e22))) & (~ (e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & (~ (e20 = op2(e21, e21)) | ~ (e21 = op2(e21, e20))) & (e21 = op2(e21, e21)) & (e21 = op2(e21, e21))) | ((~ (op2(e20, e20) = e23) | ~ (e20 = op2(e20, e23))) & (~ (op2(e20, e20) = e22) | ~ (e20 = op2(e20, e22))) & (~ (op2(e20, e20) = e21) | ~ (e20 = op2(e20, e21))) & (~ (e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20))) & (e21 = op2(e20, e21)) & (e21 = op2(e21, e20))) | ((~ (e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23))) & (~ (e22 = op2(e23, e23)) | ~ (e23 = op2(e23, e22))) & (~ (e21 = op2(e23, e23)) | ~ (e23 = op2(e23, e21))) & (~ (e20 = op2(e23, e23)) | ~ (e23 = op2(e23, e20))) & (e20 = op2(e23, e20)) & (e20 = op2(e20, e23))) | ((~ (e23 = op2(e22, e22)) | ~ (e22 = op2(e22, e23))) & (~ (e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & (~ (e21 = op2(e22, e22)) | ~ (e22 = op2(e22, e21))) & (~ (e20 = op2(e22, e22)) | ~ (e22 = op2(e22, e20))) & (e20 = op2(e22, e20)) & (e20 = op2(e20, e22))) | ((~ (e23 = op2(e21, e21)) | ~ (e21 = op2(e21, e23))) & (~ (e22 = op2(e21, e21)) | ~ (e21 = op2(e21, e22))) & (~ (e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & (~ (e20 = op2(e21, e21)) | ~ (e21 = op2(e21, e20))) & (e20 = op2(e21, e20)) & (e20 = op2(e20, e21))) | ((~ (op2(e20, e20) = e23) | ~ (e20 = op2(e20, e23))) & (~ (op2(e20, e20) = e22) | ~ (e20 = op2(e20, e22))) & (~ (op2(e20, e20) = e21) | ~ (e20 = op2(e20, e21))) & (~ (e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20))) & (e20 = op2(e20, e20)) & (e20 = op2(e20, e20)))) & ((((e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23))) & ((e22 = op2(e22, e23)) | ~ (e23 = op2(e22, e22))) & ((e21 = op2(e21, e23)) | ~ (e23 = op2(e21, e21))) & ((e20 = op2(e20, e23)) | ~ (op2(e20, e20) = e23)) & (e23 = op2(e23, op2(e23, e23))) & (e22 = op2(e23, op2(e23, e22))) & (e21 = op2(e23, op2(e23, e21))) & (e20 = op2(e23, op2(e23, e20)))) | (((e23 = op2(e23, e22)) | ~ (e22 = op2(e23, e23))) & ((e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & ((e21 = op2(e21, e22)) | ~ (e22 = op2(e21, e21))) & ((e20 = op2(e20, e22)) | ~ (op2(e20, e20) = e22)) & (e23 = op2(e22, op2(e22, e23))) & (e22 = op2(e22, op2(e22, e22))) & (e21 = op2(e22, op2(e22, e21))) & (e20 = op2(e22, op2(e22, e20)))) | (((e23 = op2(e23, e21)) | ~ (e21 = op2(e23, e23))) & ((e22 = op2(e22, e21)) | ~ (e21 = op2(e22, e22))) & ((e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & ((e20 = op2(e20, e21)) | ~ (op2(e20, e20) = e21)) & (e23 = op2(e21, op2(e21, e23))) & (e22 = op2(e21, op2(e21, e22))) & (e21 = op2(e21, op2(e21, e21))) & (e20 = op2(e21, op2(e21, e20)))) | (((e23 = op2(e23, e20)) | ~ (e20 = op2(e23, e23))) & ((e22 = op2(e22, e20)) | ~ (e20 = op2(e22, e22))) & ((e21 = op2(e21, e20)) | ~ (e20 = op2(e21, e21))) & ((e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20))) & (e23 = op2(e20, op2(e20, e23))) & (e22 = op2(e20, op2(e20, e22))) & (e21 = op2(e20, op2(e20, e21))) & (e20 = op2(e20, op2(e20, e20)))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG126+1.p', ax11)).
fof(f1818, plain, (spl48_177 | spl48_176 | spl48_175 | spl48_174 | spl48_173 | spl48_172 | spl48_171 | spl48_170 | spl48_169 | spl48_168 | spl48_167 | spl48_166 | spl48_165 | spl48_164 | spl48_163 | spl48_68), inference(avatar_split_clause, [], [f590, f1031, f1583, f1593, f1603, f1613, f1623, f1633, f1643, f1653, f1663, f1673, f1683, f1693, f1703, f1713, f1723])).
fof(f1723, plain, (spl48_177 <=> sP21), introduced(avatar_definition, [new_symbols(naming, [spl48_177])])).
fof(f1713, plain, (spl48_176 <=> sP22), introduced(avatar_definition, [new_symbols(naming, [spl48_176])])).
fof(f1703, plain, (spl48_175 <=> sP23), introduced(avatar_definition, [new_symbols(naming, [spl48_175])])).
fof(f1693, plain, (spl48_174 <=> sP24), introduced(avatar_definition, [new_symbols(naming, [spl48_174])])).
fof(f1683, plain, (spl48_173 <=> sP25), introduced(avatar_definition, [new_symbols(naming, [spl48_173])])).
fof(f1673, plain, (spl48_172 <=> sP26), introduced(avatar_definition, [new_symbols(naming, [spl48_172])])).
fof(f1663, plain, (spl48_171 <=> sP27), introduced(avatar_definition, [new_symbols(naming, [spl48_171])])).
fof(f1653, plain, (spl48_170 <=> sP28), introduced(avatar_definition, [new_symbols(naming, [spl48_170])])).
fof(f1643, plain, (spl48_169 <=> sP29), introduced(avatar_definition, [new_symbols(naming, [spl48_169])])).
fof(f1633, plain, (spl48_168 <=> sP30), introduced(avatar_definition, [new_symbols(naming, [spl48_168])])).
fof(f1623, plain, (spl48_167 <=> sP31), introduced(avatar_definition, [new_symbols(naming, [spl48_167])])).
fof(f1613, plain, (spl48_166 <=> sP32), introduced(avatar_definition, [new_symbols(naming, [spl48_166])])).
fof(f1603, plain, (spl48_165 <=> sP33), introduced(avatar_definition, [new_symbols(naming, [spl48_165])])).
fof(f1593, plain, (spl48_164 <=> sP34), introduced(avatar_definition, [new_symbols(naming, [spl48_164])])).
fof(f1583, plain, (spl48_163 <=> sP35), introduced(avatar_definition, [new_symbols(naming, [spl48_163])])).
fof(f1031, plain, (spl48_68 <=> (e23 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl48_68])])).
fof(f590, plain, ((e23 = op2(e23, e23)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21), inference(cnf_transformation, [], [f58])).
fof(f1813, plain, (spl48_177 | spl48_176 | spl48_175 | spl48_174 | spl48_173 | spl48_172 | spl48_171 | spl48_170 | spl48_169 | spl48_168 | spl48_167 | spl48_166 | spl48_165 | spl48_164 | spl48_163 | ~ spl48_68), inference(avatar_split_clause, [], [f682, f1031, f1583, f1593, f1603, f1613, f1623, f1633, f1643, f1653, f1663, f1673, f1683, f1693, f1703, f1713, f1723])).
fof(f682, plain, (~ (e23 = op2(e23, e23)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21), inference(duplicate_literal_removal, [], [f595])).
fof(f595, plain, (~ (e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21), inference(cnf_transformation, [], [f58])).
fof(f1790, plain, (~ spl48_188 | ~ spl48_65 | spl48_80), inference(avatar_split_clause, [], [f581, f1082, f1019, f1787])).
fof(f581, plain, ((e23 = op2(e23, e20)) | ~ (e20 = op2(e23, e23)) | ~ sP18), inference(cnf_transformation, [], [f107])).
fof(f107, plain, ((((e23 = op2(e23, e20)) | ~ (e20 = op2(e23, e23))) & ((e22 = op2(e22, e20)) | ~ (e20 = op2(e22, e22))) & ((e21 = op2(e21, e20)) | ~ (e20 = op2(e21, e21))) & ((e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20))) & (e23 = op2(e20, op2(e20, e23))) & (e22 = op2(e20, op2(e20, e22))) & (e21 = op2(e20, op2(e20, e21))) & (e20 = op2(e20, op2(e20, e20)))) | ~ sP18), inference(nnf_transformation, [], [f40])).
fof(f1764, plain, (~ spl48_183 | ~ spl48_86 | spl48_91), inference(avatar_split_clause, [], [f572, f1129, f1108, f1760])).
fof(f572, plain, ((e22 = op2(e22, e21)) | ~ (e21 = op2(e22, e22)) | ~ sP19), inference(cnf_transformation, [], [f106])).
fof(f106, plain, ((((e23 = op2(e23, e21)) | ~ (e21 = op2(e23, e23))) & ((e22 = op2(e22, e21)) | ~ (e21 = op2(e22, e22))) & ((e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & ((e20 = op2(e20, e21)) | ~ (op2(e20, e20) = e21)) & (e23 = op2(e21, op2(e21, e23))) & (e22 = op2(e21, op2(e21, e22))) & (e21 = op2(e21, op2(e21, e21))) & (e20 = op2(e21, op2(e21, e20)))) | ~ sP19), inference(nnf_transformation, [], [f41])).
fof(f1748, plain, (~ spl48_178 | spl48_180), inference(avatar_split_clause, [], [f560, f1745, f1733])).
fof(f560, plain, ((e22 = op2(e22, op2(e22, e22))) | ~ sP20), inference(cnf_transformation, [], [f105])).
fof(f105, plain, ((((e23 = op2(e23, e22)) | ~ (e22 = op2(e23, e23))) & ((e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & ((e21 = op2(e21, e22)) | ~ (e22 = op2(e21, e21))) & ((e20 = op2(e20, e22)) | ~ (op2(e20, e20) = e22)) & (e23 = op2(e22, op2(e22, e23))) & (e22 = op2(e22, op2(e22, e22))) & (e21 = op2(e22, op2(e22, e21))) & (e20 = op2(e22, op2(e22, e20)))) | ~ sP20), inference(nnf_transformation, [], [f42])).
fof(f1731, plain, (~ spl48_177 | spl48_125), inference(avatar_split_clause, [], [f552, f1274, f1723])).
fof(f552, plain, ((e20 = op2(e20, e20)) | ~ sP21), inference(cnf_transformation, [], [f104])).
fof(f104, plain, (((~ (op2(e20, e20) = e23) | ~ (e20 = op2(e20, e23))) & (~ (op2(e20, e20) = e22) | ~ (e20 = op2(e20, e22))) & (~ (op2(e20, e20) = e21) | ~ (e20 = op2(e20, e21))) & (~ (e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20))) & (e20 = op2(e20, e20)) & (e20 = op2(e20, e20))) | ~ sP21), inference(nnf_transformation, [], [f43])).
fof(f1721, plain, (~ spl48_176 | spl48_121), inference(avatar_split_clause, [], [f546, f1257, f1713])).
fof(f546, plain, ((e20 = op2(e20, e21)) | ~ sP22), inference(cnf_transformation, [], [f103])).
fof(f103, plain, (((~ (e23 = op2(e21, e21)) | ~ (e21 = op2(e21, e23))) & (~ (e22 = op2(e21, e21)) | ~ (e21 = op2(e21, e22))) & (~ (e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & (~ (e20 = op2(e21, e21)) | ~ (e21 = op2(e21, e20))) & (e20 = op2(e21, e20)) & (e20 = op2(e20, e21))) | ~ sP22), inference(nnf_transformation, [], [f44])).
fof(f1711, plain, (~ spl48_175 | spl48_117), inference(avatar_split_clause, [], [f540, f1240, f1703])).
fof(f540, plain, ((e20 = op2(e20, e22)) | ~ sP23), inference(cnf_transformation, [], [f102])).
fof(f102, plain, (((~ (e23 = op2(e22, e22)) | ~ (e22 = op2(e22, e23))) & (~ (e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & (~ (e21 = op2(e22, e22)) | ~ (e22 = op2(e22, e21))) & (~ (e20 = op2(e22, e22)) | ~ (e22 = op2(e22, e20))) & (e20 = op2(e22, e20)) & (e20 = op2(e20, e22))) | ~ sP23), inference(nnf_transformation, [], [f45])).
fof(f1707, plain, (~ spl48_175 | ~ spl48_87), inference(avatar_split_clause, [], [f685, f1112, f1703])).
fof(f685, plain, (~ (e22 = op2(e22, e22)) | ~ sP23), inference(duplicate_literal_removal, [], [f544])).
fof(f544, plain, (~ (e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22)) | ~ sP23), inference(cnf_transformation, [], [f102])).
fof(f1706, plain, (~ spl48_175 | ~ spl48_83 | ~ spl48_88), inference(avatar_split_clause, [], [f545, f1116, f1095, f1703])).
fof(f545, plain, (~ (e23 = op2(e22, e22)) | ~ (e22 = op2(e22, e23)) | ~ sP23), inference(cnf_transformation, [], [f102])).
fof(f1700, plain, (~ spl48_174 | spl48_77), inference(avatar_split_clause, [], [f535, f1070, f1693])).
fof(f535, plain, ((e20 = op2(e23, e20)) | ~ sP24), inference(cnf_transformation, [], [f101])).
fof(f101, plain, (((~ (e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23))) & (~ (e22 = op2(e23, e23)) | ~ (e23 = op2(e23, e22))) & (~ (e21 = op2(e23, e23)) | ~ (e23 = op2(e23, e21))) & (~ (e20 = op2(e23, e23)) | ~ (e23 = op2(e23, e20))) & (e20 = op2(e23, e20)) & (e20 = op2(e20, e23))) | ~ sP24), inference(nnf_transformation, [], [f46])).
fof(f1690, plain, (~ spl48_173 | spl48_122), inference(avatar_split_clause, [], [f529, f1261, f1683])).
fof(f529, plain, ((e21 = op2(e20, e21)) | ~ sP25), inference(cnf_transformation, [], [f100])).
fof(f100, plain, (((~ (op2(e20, e20) = e23) | ~ (e20 = op2(e20, e23))) & (~ (op2(e20, e20) = e22) | ~ (e20 = op2(e20, e22))) & (~ (op2(e20, e20) = e21) | ~ (e20 = op2(e20, e21))) & (~ (e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20))) & (e21 = op2(e20, e21)) & (e21 = op2(e21, e20))) | ~ sP25), inference(nnf_transformation, [], [f47])).
fof(f1686, plain, (~ spl48_173 | ~ spl48_113 | ~ spl48_128), inference(avatar_split_clause, [], [f533, f1286, f1223, f1683])).
fof(f533, plain, (~ (op2(e20, e20) = e23) | ~ (e20 = op2(e20, e23)) | ~ sP25), inference(cnf_transformation, [], [f100])).
fof(f1679, plain, (~ spl48_172 | ~ spl48_110 | ~ spl48_105), inference(avatar_split_clause, [], [f524, f1189, f1210, f1673])).
fof(f524, plain, (~ (e20 = op2(e21, e21)) | ~ (e21 = op2(e21, e20)) | ~ sP26), inference(cnf_transformation, [], [f99])).
fof(f99, plain, (((~ (e23 = op2(e21, e21)) | ~ (e21 = op2(e21, e23))) & (~ (e22 = op2(e21, e21)) | ~ (e21 = op2(e21, e22))) & (~ (e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & (~ (e20 = op2(e21, e21)) | ~ (e21 = op2(e21, e20))) & (e21 = op2(e21, e21)) & (e21 = op2(e21, e21))) | ~ sP26), inference(nnf_transformation, [], [f48])).
fof(f1670, plain, (~ spl48_171 | spl48_90), inference(avatar_split_clause, [], [f517, f1125, f1663])).
fof(f517, plain, ((e21 = op2(e22, e21)) | ~ sP27), inference(cnf_transformation, [], [f98])).
fof(f98, plain, (((~ (e23 = op2(e22, e22)) | ~ (e22 = op2(e22, e23))) & (~ (e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & (~ (e21 = op2(e22, e22)) | ~ (e22 = op2(e22, e21))) & (~ (e20 = op2(e22, e22)) | ~ (e22 = op2(e22, e20))) & (e21 = op2(e22, e21)) & (e21 = op2(e21, e22))) | ~ sP27), inference(nnf_transformation, [], [f49])).
fof(f1667, plain, (~ spl48_171 | ~ spl48_87), inference(avatar_split_clause, [], [f689, f1112, f1663])).
fof(f689, plain, (~ (e22 = op2(e22, e22)) | ~ sP27), inference(duplicate_literal_removal, [], [f520])).
fof(f520, plain, (~ (e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22)) | ~ sP27), inference(cnf_transformation, [], [f98])).
fof(f1666, plain, (~ spl48_171 | ~ spl48_83 | ~ spl48_88), inference(avatar_split_clause, [], [f521, f1116, f1095, f1663])).
fof(f521, plain, (~ (e23 = op2(e22, e22)) | ~ (e22 = op2(e22, e23)) | ~ sP27), inference(cnf_transformation, [], [f98])).
fof(f1661, plain, (~ spl48_170 | spl48_98), inference(avatar_split_clause, [], [f510, f1159, f1653])).
fof(f510, plain, ((e21 = op2(e21, e23)) | ~ sP28), inference(cnf_transformation, [], [f97])).
fof(f97, plain, (((~ (e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23))) & (~ (e22 = op2(e23, e23)) | ~ (e23 = op2(e23, e22))) & (~ (e21 = op2(e23, e23)) | ~ (e23 = op2(e23, e21))) & (~ (e20 = op2(e23, e23)) | ~ (e23 = op2(e23, e20))) & (e21 = op2(e23, e21)) & (e21 = op2(e21, e23))) | ~ sP28), inference(nnf_transformation, [], [f50])).
fof(f1651, plain, (~ spl48_169 | spl48_95), inference(avatar_split_clause, [], [f504, f1146, f1643])).
fof(f504, plain, ((e22 = op2(e22, e20)) | ~ sP29), inference(cnf_transformation, [], [f96])).
fof(f96, plain, (((~ (op2(e20, e20) = e23) | ~ (e20 = op2(e20, e23))) & (~ (op2(e20, e20) = e22) | ~ (e20 = op2(e20, e22))) & (~ (op2(e20, e20) = e21) | ~ (e20 = op2(e20, e21))) & (~ (e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20))) & (e22 = op2(e20, e22)) & (e22 = op2(e22, e20))) | ~ sP29), inference(nnf_transformation, [], [f51])).
fof(f1639, plain, (~ spl48_168 | ~ spl48_110 | ~ spl48_105), inference(avatar_split_clause, [], [f500, f1189, f1210, f1633])).
fof(f500, plain, (~ (e20 = op2(e21, e21)) | ~ (e21 = op2(e21, e20)) | ~ sP30), inference(cnf_transformation, [], [f95])).
fof(f95, plain, (((~ (e23 = op2(e21, e21)) | ~ (e21 = op2(e21, e23))) & (~ (e22 = op2(e21, e21)) | ~ (e21 = op2(e21, e22))) & (~ (e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & (~ (e20 = op2(e21, e21)) | ~ (e21 = op2(e21, e20))) & (e22 = op2(e21, e22)) & (e22 = op2(e22, e21))) | ~ sP30), inference(nnf_transformation, [], [f52])).
fof(f1631, plain, (~ spl48_167 | spl48_87), inference(avatar_split_clause, [], [f492, f1112, f1623])).
fof(f492, plain, ((e22 = op2(e22, e22)) | ~ sP31), inference(cnf_transformation, [], [f94])).
fof(f94, plain, (((~ (e23 = op2(e22, e22)) | ~ (e22 = op2(e22, e23))) & (~ (e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & (~ (e21 = op2(e22, e22)) | ~ (e22 = op2(e22, e21))) & (~ (e20 = op2(e22, e22)) | ~ (e22 = op2(e22, e20))) & (e22 = op2(e22, e22)) & (e22 = op2(e22, e22))) | ~ sP31), inference(nnf_transformation, [], [f53])).
fof(f1627, plain, (~ spl48_167 | ~ spl48_87), inference(avatar_split_clause, [], [f693, f1112, f1623])).
fof(f693, plain, (~ (e22 = op2(e22, e22)) | ~ sP31), inference(duplicate_literal_removal, [], [f496])).
fof(f496, plain, (~ (e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22)) | ~ sP31), inference(cnf_transformation, [], [f94])).
fof(f1620, plain, (~ spl48_166 | spl48_71), inference(avatar_split_clause, [], [f487, f1044, f1613])).
fof(f487, plain, ((e22 = op2(e23, e22)) | ~ sP32), inference(cnf_transformation, [], [f93])).
fof(f93, plain, (((~ (e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23))) & (~ (e22 = op2(e23, e23)) | ~ (e23 = op2(e23, e22))) & (~ (e21 = op2(e23, e23)) | ~ (e23 = op2(e23, e21))) & (~ (e20 = op2(e23, e23)) | ~ (e23 = op2(e23, e20))) & (e22 = op2(e23, e22)) & (e22 = op2(e22, e23))) | ~ sP32), inference(nnf_transformation, [], [f54])).
fof(f1611, plain, (~ spl48_165 | spl48_80), inference(avatar_split_clause, [], [f480, f1082, f1603])).
fof(f480, plain, ((e23 = op2(e23, e20)) | ~ sP33), inference(cnf_transformation, [], [f92])).
fof(f92, plain, (((~ (op2(e20, e20) = e23) | ~ (e20 = op2(e20, e23))) & (~ (op2(e20, e20) = e22) | ~ (e20 = op2(e20, e22))) & (~ (op2(e20, e20) = e21) | ~ (e20 = op2(e20, e21))) & (~ (e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20))) & (e23 = op2(e20, e23)) & (e23 = op2(e23, e20))) | ~ sP33), inference(nnf_transformation, [], [f55])).
fof(f1599, plain, (~ spl48_164 | ~ spl48_110 | ~ spl48_105), inference(avatar_split_clause, [], [f476, f1189, f1210, f1593])).
fof(f476, plain, (~ (e20 = op2(e21, e21)) | ~ (e21 = op2(e21, e20)) | ~ sP34), inference(cnf_transformation, [], [f91])).
fof(f91, plain, (((~ (e23 = op2(e21, e21)) | ~ (e21 = op2(e21, e23))) & (~ (e22 = op2(e21, e21)) | ~ (e21 = op2(e21, e22))) & (~ (e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & (~ (e20 = op2(e21, e21)) | ~ (e21 = op2(e21, e20))) & (e23 = op2(e21, e23)) & (e23 = op2(e23, e21))) | ~ sP34), inference(nnf_transformation, [], [f56])).
fof(f1591, plain, (~ spl48_163 | spl48_72), inference(avatar_split_clause, [], [f468, f1048, f1583])).
fof(f468, plain, ((e23 = op2(e23, e22)) | ~ sP35), inference(cnf_transformation, [], [f90])).
fof(f90, plain, (((~ (e23 = op2(e22, e22)) | ~ (e22 = op2(e22, e23))) & (~ (e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & (~ (e21 = op2(e22, e22)) | ~ (e22 = op2(e22, e21))) & (~ (e20 = op2(e22, e22)) | ~ (e22 = op2(e22, e20))) & (e23 = op2(e22, e23)) & (e23 = op2(e23, e22))) | ~ sP35), inference(nnf_transformation, [], [f57])).
fof(f1590, plain, (~ spl48_163 | spl48_84), inference(avatar_split_clause, [], [f469, f1099, f1583])).
fof(f469, plain, ((e23 = op2(e22, e23)) | ~ sP35), inference(cnf_transformation, [], [f90])).
fof(f1587, plain, (~ spl48_163 | ~ spl48_87), inference(avatar_split_clause, [], [f697, f1112, f1583])).
fof(f697, plain, (~ (e22 = op2(e22, e22)) | ~ sP35), inference(duplicate_literal_removal, [], [f472])).
fof(f472, plain, (~ (e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22)) | ~ sP35), inference(cnf_transformation, [], [f90])).
fof(f1561, plain, (spl48_154 | spl48_149 | spl48_144 | ~ spl48_64 | spl48_49), inference(avatar_split_clause, [], [f458, f919, f982, f1473, f1500, f1527])).
fof(f1527, plain, (spl48_154 <=> sP0), introduced(avatar_definition, [new_symbols(naming, [spl48_154])])).
fof(f1500, plain, (spl48_149 <=> sP1), introduced(avatar_definition, [new_symbols(naming, [spl48_149])])).
fof(f1473, plain, (spl48_144 <=> sP2), introduced(avatar_definition, [new_symbols(naming, [spl48_144])])).
fof(f458, plain, ((e10 = op1(e10, e13)) | ~ (op1(e10, e10) = e13) | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f39])).
fof(f39, plain, ((((~ (e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13))) & (~ (e12 = op1(e13, e13)) | ~ (e13 = op1(e13, e12))) & (~ (e11 = op1(e13, e13)) | ~ (e13 = op1(e13, e11))) & (~ (e10 = op1(e13, e13)) | ~ (e13 = op1(e13, e10))) & (e13 = op1(e13, e13)) & (e13 = op1(e13, e13))) | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3) & ((((e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13))) & ((e12 = op1(e12, e13)) | ~ (e13 = op1(e12, e12))) & ((e11 = op1(e11, e13)) | ~ (e13 = op1(e11, e11))) & ((e10 = op1(e10, e13)) | ~ (op1(e10, e10) = e13)) & (e13 = op1(e13, op1(e13, e13))) & (e12 = op1(e13, op1(e13, e12))) & (e11 = op1(e13, op1(e13, e11))) & (e10 = op1(e13, op1(e13, e10)))) | sP2 | sP1 | sP0)), inference(definition_folding, [], [f10, e38, e37, e36, e35, e34, e33, e32, e31, e30, e29, e28, e27, e26, e25, e24, e23, e22, e21])).
fof(f21, plain, ((((e13 = op1(e13, e10)) | ~ (e10 = op1(e13, e13))) & ((e12 = op1(e12, e10)) | ~ (e10 = op1(e12, e12))) & ((e11 = op1(e11, e10)) | ~ (e10 = op1(e11, e11))) & ((e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10))) & (e13 = op1(e10, op1(e10, e13))) & (e12 = op1(e10, op1(e10, e12))) & (e11 = op1(e10, op1(e10, e11))) & (e10 = op1(e10, op1(e10, e10)))) | ~ sP0), inference(usedef, [], [e21])).
fof(e21, plain, (sP0 <=> (((e13 = op1(e13, e10)) | ~ (e10 = op1(e13, e13))) & ((e12 = op1(e12, e10)) | ~ (e10 = op1(e12, e12))) & ((e11 = op1(e11, e10)) | ~ (e10 = op1(e11, e11))) & ((e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10))) & (e13 = op1(e10, op1(e10, e13))) & (e12 = op1(e10, op1(e10, e12))) & (e11 = op1(e10, op1(e10, e11))) & (e10 = op1(e10, op1(e10, e10))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f22, plain, ((((e13 = op1(e13, e11)) | ~ (e11 = op1(e13, e13))) & ((e12 = op1(e12, e11)) | ~ (e11 = op1(e12, e12))) & ((e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & ((e10 = op1(e10, e11)) | ~ (op1(e10, e10) = e11)) & (e13 = op1(e11, op1(e11, e13))) & (e12 = op1(e11, op1(e11, e12))) & (e11 = op1(e11, op1(e11, e11))) & (e10 = op1(e11, op1(e11, e10)))) | ~ sP1), inference(usedef, [], [e22])).
fof(e22, plain, (sP1 <=> (((e13 = op1(e13, e11)) | ~ (e11 = op1(e13, e13))) & ((e12 = op1(e12, e11)) | ~ (e11 = op1(e12, e12))) & ((e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & ((e10 = op1(e10, e11)) | ~ (op1(e10, e10) = e11)) & (e13 = op1(e11, op1(e11, e13))) & (e12 = op1(e11, op1(e11, e12))) & (e11 = op1(e11, op1(e11, e11))) & (e10 = op1(e11, op1(e11, e10))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f23, plain, ((((e13 = op1(e13, e12)) | ~ (e12 = op1(e13, e13))) & ((e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & ((e11 = op1(e11, e12)) | ~ (e12 = op1(e11, e11))) & ((e10 = op1(e10, e12)) | ~ (op1(e10, e10) = e12)) & (e13 = op1(e12, op1(e12, e13))) & (e12 = op1(e12, op1(e12, e12))) & (e11 = op1(e12, op1(e12, e11))) & (e10 = op1(e12, op1(e12, e10)))) | ~ sP2), inference(usedef, [], [e23])).
fof(e23, plain, (sP2 <=> (((e13 = op1(e13, e12)) | ~ (e12 = op1(e13, e13))) & ((e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & ((e11 = op1(e11, e12)) | ~ (e12 = op1(e11, e11))) & ((e10 = op1(e10, e12)) | ~ (op1(e10, e10) = e12)) & (e13 = op1(e12, op1(e12, e13))) & (e12 = op1(e12, op1(e12, e12))) & (e11 = op1(e12, op1(e12, e11))) & (e10 = op1(e12, op1(e12, e10))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f24, plain, (((~ (op1(e10, e10) = e13) | ~ (e10 = op1(e10, e13))) & (~ (op1(e10, e10) = e12) | ~ (e10 = op1(e10, e12))) & (~ (op1(e10, e10) = e11) | ~ (e10 = op1(e10, e11))) & (~ (e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10))) & (e10 = op1(e10, e10)) & (e10 = op1(e10, e10))) | ~ sP3), inference(usedef, [], [e24])).
fof(e24, plain, (sP3 <=> ((~ (op1(e10, e10) = e13) | ~ (e10 = op1(e10, e13))) & (~ (op1(e10, e10) = e12) | ~ (e10 = op1(e10, e12))) & (~ (op1(e10, e10) = e11) | ~ (e10 = op1(e10, e11))) & (~ (e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10))) & (e10 = op1(e10, e10)) & (e10 = op1(e10, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f25, plain, (((~ (e13 = op1(e11, e11)) | ~ (e11 = op1(e11, e13))) & (~ (e12 = op1(e11, e11)) | ~ (e11 = op1(e11, e12))) & (~ (e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & (~ (e10 = op1(e11, e11)) | ~ (e11 = op1(e11, e10))) & (e10 = op1(e11, e10)) & (e10 = op1(e10, e11))) | ~ sP4), inference(usedef, [], [e25])).
fof(e25, plain, (sP4 <=> ((~ (e13 = op1(e11, e11)) | ~ (e11 = op1(e11, e13))) & (~ (e12 = op1(e11, e11)) | ~ (e11 = op1(e11, e12))) & (~ (e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & (~ (e10 = op1(e11, e11)) | ~ (e11 = op1(e11, e10))) & (e10 = op1(e11, e10)) & (e10 = op1(e10, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f26, plain, (((~ (e13 = op1(e12, e12)) | ~ (e12 = op1(e12, e13))) & (~ (e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & (~ (e11 = op1(e12, e12)) | ~ (e12 = op1(e12, e11))) & (~ (e10 = op1(e12, e12)) | ~ (e12 = op1(e12, e10))) & (e10 = op1(e12, e10)) & (e10 = op1(e10, e12))) | ~ sP5), inference(usedef, [], [e26])).
fof(e26, plain, (sP5 <=> ((~ (e13 = op1(e12, e12)) | ~ (e12 = op1(e12, e13))) & (~ (e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & (~ (e11 = op1(e12, e12)) | ~ (e12 = op1(e12, e11))) & (~ (e10 = op1(e12, e12)) | ~ (e12 = op1(e12, e10))) & (e10 = op1(e12, e10)) & (e10 = op1(e10, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f27, plain, (((~ (e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13))) & (~ (e12 = op1(e13, e13)) | ~ (e13 = op1(e13, e12))) & (~ (e11 = op1(e13, e13)) | ~ (e13 = op1(e13, e11))) & (~ (e10 = op1(e13, e13)) | ~ (e13 = op1(e13, e10))) & (e10 = op1(e13, e10)) & (e10 = op1(e10, e13))) | ~ sP6), inference(usedef, [], [e27])).
fof(e27, plain, (sP6 <=> ((~ (e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13))) & (~ (e12 = op1(e13, e13)) | ~ (e13 = op1(e13, e12))) & (~ (e11 = op1(e13, e13)) | ~ (e13 = op1(e13, e11))) & (~ (e10 = op1(e13, e13)) | ~ (e13 = op1(e13, e10))) & (e10 = op1(e13, e10)) & (e10 = op1(e10, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f28, plain, (((~ (op1(e10, e10) = e13) | ~ (e10 = op1(e10, e13))) & (~ (op1(e10, e10) = e12) | ~ (e10 = op1(e10, e12))) & (~ (op1(e10, e10) = e11) | ~ (e10 = op1(e10, e11))) & (~ (e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10))) & (e11 = op1(e10, e11)) & (e11 = op1(e11, e10))) | ~ sP7), inference(usedef, [], [e28])).
fof(e28, plain, (sP7 <=> ((~ (op1(e10, e10) = e13) | ~ (e10 = op1(e10, e13))) & (~ (op1(e10, e10) = e12) | ~ (e10 = op1(e10, e12))) & (~ (op1(e10, e10) = e11) | ~ (e10 = op1(e10, e11))) & (~ (e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10))) & (e11 = op1(e10, e11)) & (e11 = op1(e11, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f29, plain, (((~ (e13 = op1(e11, e11)) | ~ (e11 = op1(e11, e13))) & (~ (e12 = op1(e11, e11)) | ~ (e11 = op1(e11, e12))) & (~ (e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & (~ (e10 = op1(e11, e11)) | ~ (e11 = op1(e11, e10))) & (e11 = op1(e11, e11)) & (e11 = op1(e11, e11))) | ~ sP8), inference(usedef, [], [e29])).
fof(e29, plain, (sP8 <=> ((~ (e13 = op1(e11, e11)) | ~ (e11 = op1(e11, e13))) & (~ (e12 = op1(e11, e11)) | ~ (e11 = op1(e11, e12))) & (~ (e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & (~ (e10 = op1(e11, e11)) | ~ (e11 = op1(e11, e10))) & (e11 = op1(e11, e11)) & (e11 = op1(e11, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f30, plain, (((~ (e13 = op1(e12, e12)) | ~ (e12 = op1(e12, e13))) & (~ (e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & (~ (e11 = op1(e12, e12)) | ~ (e12 = op1(e12, e11))) & (~ (e10 = op1(e12, e12)) | ~ (e12 = op1(e12, e10))) & (e11 = op1(e12, e11)) & (e11 = op1(e11, e12))) | ~ sP9), inference(usedef, [], [e30])).
fof(e30, plain, (sP9 <=> ((~ (e13 = op1(e12, e12)) | ~ (e12 = op1(e12, e13))) & (~ (e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & (~ (e11 = op1(e12, e12)) | ~ (e12 = op1(e12, e11))) & (~ (e10 = op1(e12, e12)) | ~ (e12 = op1(e12, e10))) & (e11 = op1(e12, e11)) & (e11 = op1(e11, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f31, plain, (((~ (e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13))) & (~ (e12 = op1(e13, e13)) | ~ (e13 = op1(e13, e12))) & (~ (e11 = op1(e13, e13)) | ~ (e13 = op1(e13, e11))) & (~ (e10 = op1(e13, e13)) | ~ (e13 = op1(e13, e10))) & (e11 = op1(e13, e11)) & (e11 = op1(e11, e13))) | ~ sP10), inference(usedef, [], [e31])).
fof(e31, plain, (sP10 <=> ((~ (e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13))) & (~ (e12 = op1(e13, e13)) | ~ (e13 = op1(e13, e12))) & (~ (e11 = op1(e13, e13)) | ~ (e13 = op1(e13, e11))) & (~ (e10 = op1(e13, e13)) | ~ (e13 = op1(e13, e10))) & (e11 = op1(e13, e11)) & (e11 = op1(e11, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f32, plain, (((~ (op1(e10, e10) = e13) | ~ (e10 = op1(e10, e13))) & (~ (op1(e10, e10) = e12) | ~ (e10 = op1(e10, e12))) & (~ (op1(e10, e10) = e11) | ~ (e10 = op1(e10, e11))) & (~ (e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10))) & (e12 = op1(e10, e12)) & (e12 = op1(e12, e10))) | ~ sP11), inference(usedef, [], [e32])).
fof(e32, plain, (sP11 <=> ((~ (op1(e10, e10) = e13) | ~ (e10 = op1(e10, e13))) & (~ (op1(e10, e10) = e12) | ~ (e10 = op1(e10, e12))) & (~ (op1(e10, e10) = e11) | ~ (e10 = op1(e10, e11))) & (~ (e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10))) & (e12 = op1(e10, e12)) & (e12 = op1(e12, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f33, plain, (((~ (e13 = op1(e11, e11)) | ~ (e11 = op1(e11, e13))) & (~ (e12 = op1(e11, e11)) | ~ (e11 = op1(e11, e12))) & (~ (e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & (~ (e10 = op1(e11, e11)) | ~ (e11 = op1(e11, e10))) & (e12 = op1(e11, e12)) & (e12 = op1(e12, e11))) | ~ sP12), inference(usedef, [], [e33])).
fof(e33, plain, (sP12 <=> ((~ (e13 = op1(e11, e11)) | ~ (e11 = op1(e11, e13))) & (~ (e12 = op1(e11, e11)) | ~ (e11 = op1(e11, e12))) & (~ (e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & (~ (e10 = op1(e11, e11)) | ~ (e11 = op1(e11, e10))) & (e12 = op1(e11, e12)) & (e12 = op1(e12, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f34, plain, (((~ (e13 = op1(e12, e12)) | ~ (e12 = op1(e12, e13))) & (~ (e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & (~ (e11 = op1(e12, e12)) | ~ (e12 = op1(e12, e11))) & (~ (e10 = op1(e12, e12)) | ~ (e12 = op1(e12, e10))) & (e12 = op1(e12, e12)) & (e12 = op1(e12, e12))) | ~ sP13), inference(usedef, [], [e34])).
fof(e34, plain, (sP13 <=> ((~ (e13 = op1(e12, e12)) | ~ (e12 = op1(e12, e13))) & (~ (e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & (~ (e11 = op1(e12, e12)) | ~ (e12 = op1(e12, e11))) & (~ (e10 = op1(e12, e12)) | ~ (e12 = op1(e12, e10))) & (e12 = op1(e12, e12)) & (e12 = op1(e12, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f35, plain, (((~ (e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13))) & (~ (e12 = op1(e13, e13)) | ~ (e13 = op1(e13, e12))) & (~ (e11 = op1(e13, e13)) | ~ (e13 = op1(e13, e11))) & (~ (e10 = op1(e13, e13)) | ~ (e13 = op1(e13, e10))) & (e12 = op1(e13, e12)) & (e12 = op1(e12, e13))) | ~ sP14), inference(usedef, [], [e35])).
fof(e35, plain, (sP14 <=> ((~ (e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13))) & (~ (e12 = op1(e13, e13)) | ~ (e13 = op1(e13, e12))) & (~ (e11 = op1(e13, e13)) | ~ (e13 = op1(e13, e11))) & (~ (e10 = op1(e13, e13)) | ~ (e13 = op1(e13, e10))) & (e12 = op1(e13, e12)) & (e12 = op1(e12, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f36, plain, (((~ (op1(e10, e10) = e13) | ~ (e10 = op1(e10, e13))) & (~ (op1(e10, e10) = e12) | ~ (e10 = op1(e10, e12))) & (~ (op1(e10, e10) = e11) | ~ (e10 = op1(e10, e11))) & (~ (e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10))) & (e13 = op1(e10, e13)) & (e13 = op1(e13, e10))) | ~ sP15), inference(usedef, [], [e36])).
fof(e36, plain, (sP15 <=> ((~ (op1(e10, e10) = e13) | ~ (e10 = op1(e10, e13))) & (~ (op1(e10, e10) = e12) | ~ (e10 = op1(e10, e12))) & (~ (op1(e10, e10) = e11) | ~ (e10 = op1(e10, e11))) & (~ (e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10))) & (e13 = op1(e10, e13)) & (e13 = op1(e13, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f37, plain, (((~ (e13 = op1(e11, e11)) | ~ (e11 = op1(e11, e13))) & (~ (e12 = op1(e11, e11)) | ~ (e11 = op1(e11, e12))) & (~ (e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & (~ (e10 = op1(e11, e11)) | ~ (e11 = op1(e11, e10))) & (e13 = op1(e11, e13)) & (e13 = op1(e13, e11))) | ~ sP16), inference(usedef, [], [e37])).
fof(e37, plain, (sP16 <=> ((~ (e13 = op1(e11, e11)) | ~ (e11 = op1(e11, e13))) & (~ (e12 = op1(e11, e11)) | ~ (e11 = op1(e11, e12))) & (~ (e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & (~ (e10 = op1(e11, e11)) | ~ (e11 = op1(e11, e10))) & (e13 = op1(e11, e13)) & (e13 = op1(e13, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f38, plain, (((~ (e13 = op1(e12, e12)) | ~ (e12 = op1(e12, e13))) & (~ (e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & (~ (e11 = op1(e12, e12)) | ~ (e12 = op1(e12, e11))) & (~ (e10 = op1(e12, e12)) | ~ (e12 = op1(e12, e10))) & (e13 = op1(e12, e13)) & (e13 = op1(e13, e12))) | ~ sP17), inference(usedef, [], [e38])).
fof(e38, plain, (sP17 <=> ((~ (e13 = op1(e12, e12)) | ~ (e12 = op1(e12, e13))) & (~ (e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & (~ (e11 = op1(e12, e12)) | ~ (e12 = op1(e12, e11))) & (~ (e10 = op1(e12, e12)) | ~ (e12 = op1(e12, e10))) & (e13 = op1(e12, e13)) & (e13 = op1(e13, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f10, plain, ((((~ (e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13))) & (~ (e12 = op1(e13, e13)) | ~ (e13 = op1(e13, e12))) & (~ (e11 = op1(e13, e13)) | ~ (e13 = op1(e13, e11))) & (~ (e10 = op1(e13, e13)) | ~ (e13 = op1(e13, e10))) & (e13 = op1(e13, e13)) & (e13 = op1(e13, e13))) | ((~ (e13 = op1(e12, e12)) | ~ (e12 = op1(e12, e13))) & (~ (e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & (~ (e11 = op1(e12, e12)) | ~ (e12 = op1(e12, e11))) & (~ (e10 = op1(e12, e12)) | ~ (e12 = op1(e12, e10))) & (e13 = op1(e12, e13)) & (e13 = op1(e13, e12))) | ((~ (e13 = op1(e11, e11)) | ~ (e11 = op1(e11, e13))) & (~ (e12 = op1(e11, e11)) | ~ (e11 = op1(e11, e12))) & (~ (e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & (~ (e10 = op1(e11, e11)) | ~ (e11 = op1(e11, e10))) & (e13 = op1(e11, e13)) & (e13 = op1(e13, e11))) | ((~ (op1(e10, e10) = e13) | ~ (e10 = op1(e10, e13))) & (~ (op1(e10, e10) = e12) | ~ (e10 = op1(e10, e12))) & (~ (op1(e10, e10) = e11) | ~ (e10 = op1(e10, e11))) & (~ (e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10))) & (e13 = op1(e10, e13)) & (e13 = op1(e13, e10))) | ((~ (e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13))) & (~ (e12 = op1(e13, e13)) | ~ (e13 = op1(e13, e12))) & (~ (e11 = op1(e13, e13)) | ~ (e13 = op1(e13, e11))) & (~ (e10 = op1(e13, e13)) | ~ (e13 = op1(e13, e10))) & (e12 = op1(e13, e12)) & (e12 = op1(e12, e13))) | ((~ (e13 = op1(e12, e12)) | ~ (e12 = op1(e12, e13))) & (~ (e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & (~ (e11 = op1(e12, e12)) | ~ (e12 = op1(e12, e11))) & (~ (e10 = op1(e12, e12)) | ~ (e12 = op1(e12, e10))) & (e12 = op1(e12, e12)) & (e12 = op1(e12, e12))) | ((~ (e13 = op1(e11, e11)) | ~ (e11 = op1(e11, e13))) & (~ (e12 = op1(e11, e11)) | ~ (e11 = op1(e11, e12))) & (~ (e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & (~ (e10 = op1(e11, e11)) | ~ (e11 = op1(e11, e10))) & (e12 = op1(e11, e12)) & (e12 = op1(e12, e11))) | ((~ (op1(e10, e10) = e13) | ~ (e10 = op1(e10, e13))) & (~ (op1(e10, e10) = e12) | ~ (e10 = op1(e10, e12))) & (~ (op1(e10, e10) = e11) | ~ (e10 = op1(e10, e11))) & (~ (e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10))) & (e12 = op1(e10, e12)) & (e12 = op1(e12, e10))) | ((~ (e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13))) & (~ (e12 = op1(e13, e13)) | ~ (e13 = op1(e13, e12))) & (~ (e11 = op1(e13, e13)) | ~ (e13 = op1(e13, e11))) & (~ (e10 = op1(e13, e13)) | ~ (e13 = op1(e13, e10))) & (e11 = op1(e13, e11)) & (e11 = op1(e11, e13))) | ((~ (e13 = op1(e12, e12)) | ~ (e12 = op1(e12, e13))) & (~ (e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & (~ (e11 = op1(e12, e12)) | ~ (e12 = op1(e12, e11))) & (~ (e10 = op1(e12, e12)) | ~ (e12 = op1(e12, e10))) & (e11 = op1(e12, e11)) & (e11 = op1(e11, e12))) | ((~ (e13 = op1(e11, e11)) | ~ (e11 = op1(e11, e13))) & (~ (e12 = op1(e11, e11)) | ~ (e11 = op1(e11, e12))) & (~ (e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & (~ (e10 = op1(e11, e11)) | ~ (e11 = op1(e11, e10))) & (e11 = op1(e11, e11)) & (e11 = op1(e11, e11))) | ((~ (op1(e10, e10) = e13) | ~ (e10 = op1(e10, e13))) & (~ (op1(e10, e10) = e12) | ~ (e10 = op1(e10, e12))) & (~ (op1(e10, e10) = e11) | ~ (e10 = op1(e10, e11))) & (~ (e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10))) & (e11 = op1(e10, e11)) & (e11 = op1(e11, e10))) | ((~ (e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13))) & (~ (e12 = op1(e13, e13)) | ~ (e13 = op1(e13, e12))) & (~ (e11 = op1(e13, e13)) | ~ (e13 = op1(e13, e11))) & (~ (e10 = op1(e13, e13)) | ~ (e13 = op1(e13, e10))) & (e10 = op1(e13, e10)) & (e10 = op1(e10, e13))) | ((~ (e13 = op1(e12, e12)) | ~ (e12 = op1(e12, e13))) & (~ (e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & (~ (e11 = op1(e12, e12)) | ~ (e12 = op1(e12, e11))) & (~ (e10 = op1(e12, e12)) | ~ (e12 = op1(e12, e10))) & (e10 = op1(e12, e10)) & (e10 = op1(e10, e12))) | ((~ (e13 = op1(e11, e11)) | ~ (e11 = op1(e11, e13))) & (~ (e12 = op1(e11, e11)) | ~ (e11 = op1(e11, e12))) & (~ (e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & (~ (e10 = op1(e11, e11)) | ~ (e11 = op1(e11, e10))) & (e10 = op1(e11, e10)) & (e10 = op1(e10, e11))) | ((~ (op1(e10, e10) = e13) | ~ (e10 = op1(e10, e13))) & (~ (op1(e10, e10) = e12) | ~ (e10 = op1(e10, e12))) & (~ (op1(e10, e10) = e11) | ~ (e10 = op1(e10, e11))) & (~ (e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10))) & (e10 = op1(e10, e10)) & (e10 = op1(e10, e10)))) & ((((e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13))) & ((e12 = op1(e12, e13)) | ~ (e13 = op1(e12, e12))) & ((e11 = op1(e11, e13)) | ~ (e13 = op1(e11, e11))) & ((e10 = op1(e10, e13)) | ~ (op1(e10, e10) = e13)) & (e13 = op1(e13, op1(e13, e13))) & (e12 = op1(e13, op1(e13, e12))) & (e11 = op1(e13, op1(e13, e11))) & (e10 = op1(e13, op1(e13, e10)))) | (((e13 = op1(e13, e12)) | ~ (e12 = op1(e13, e13))) & ((e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & ((e11 = op1(e11, e12)) | ~ (e12 = op1(e11, e11))) & ((e10 = op1(e10, e12)) | ~ (op1(e10, e10) = e12)) & (e13 = op1(e12, op1(e12, e13))) & (e12 = op1(e12, op1(e12, e12))) & (e11 = op1(e12, op1(e12, e11))) & (e10 = op1(e12, op1(e12, e10)))) | (((e13 = op1(e13, e11)) | ~ (e11 = op1(e13, e13))) & ((e12 = op1(e12, e11)) | ~ (e11 = op1(e12, e12))) & ((e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & ((e10 = op1(e10, e11)) | ~ (op1(e10, e10) = e11)) & (e13 = op1(e11, op1(e11, e13))) & (e12 = op1(e11, op1(e11, e12))) & (e11 = op1(e11, op1(e11, e11))) & (e10 = op1(e11, op1(e11, e10)))) | (((e13 = op1(e13, e10)) | ~ (e10 = op1(e13, e13))) & ((e12 = op1(e12, e10)) | ~ (e10 = op1(e12, e12))) & ((e11 = op1(e11, e10)) | ~ (e10 = op1(e11, e11))) & ((e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10))) & (e13 = op1(e10, op1(e10, e13))) & (e12 = op1(e10, op1(e10, e12))) & (e11 = op1(e10, op1(e10, e11))) & (e10 = op1(e10, op1(e10, e10)))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG126+1.p', ax10)).
fof(f1558, plain, (spl48_143 | spl48_142 | spl48_141 | spl48_140 | spl48_139 | spl48_138 | spl48_137 | spl48_136 | spl48_135 | spl48_134 | spl48_133 | spl48_132 | spl48_131 | spl48_130 | spl48_129 | spl48_4), inference(avatar_split_clause, [], [f462, f727, f1323, f1333, f1343, f1353, f1363, f1373, f1383, f1393, f1403, f1413, f1423, f1433, f1443, f1453, f1463])).
fof(f1463, plain, (spl48_143 <=> sP3), introduced(avatar_definition, [new_symbols(naming, [spl48_143])])).
fof(f1453, plain, (spl48_142 <=> sP4), introduced(avatar_definition, [new_symbols(naming, [spl48_142])])).
fof(f1443, plain, (spl48_141 <=> sP5), introduced(avatar_definition, [new_symbols(naming, [spl48_141])])).
fof(f1433, plain, (spl48_140 <=> sP6), introduced(avatar_definition, [new_symbols(naming, [spl48_140])])).
fof(f1423, plain, (spl48_139 <=> sP7), introduced(avatar_definition, [new_symbols(naming, [spl48_139])])).
fof(f1413, plain, (spl48_138 <=> sP8), introduced(avatar_definition, [new_symbols(naming, [spl48_138])])).
fof(f1403, plain, (spl48_137 <=> sP9), introduced(avatar_definition, [new_symbols(naming, [spl48_137])])).
fof(f1393, plain, (spl48_136 <=> sP10), introduced(avatar_definition, [new_symbols(naming, [spl48_136])])).
fof(f1383, plain, (spl48_135 <=> sP11), introduced(avatar_definition, [new_symbols(naming, [spl48_135])])).
fof(f1373, plain, (spl48_134 <=> sP12), introduced(avatar_definition, [new_symbols(naming, [spl48_134])])).
fof(f1363, plain, (spl48_133 <=> sP13), introduced(avatar_definition, [new_symbols(naming, [spl48_133])])).
fof(f1353, plain, (spl48_132 <=> sP14), introduced(avatar_definition, [new_symbols(naming, [spl48_132])])).
fof(f1343, plain, (spl48_131 <=> sP15), introduced(avatar_definition, [new_symbols(naming, [spl48_131])])).
fof(f1333, plain, (spl48_130 <=> sP16), introduced(avatar_definition, [new_symbols(naming, [spl48_130])])).
fof(f1323, plain, (spl48_129 <=> sP17), introduced(avatar_definition, [new_symbols(naming, [spl48_129])])).
fof(f727, plain, (spl48_4 <=> (e13 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl48_4])])).
fof(f462, plain, ((e13 = op1(e13, e13)) | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3), inference(cnf_transformation, [], [f39])).
fof(f1553, plain, (spl48_143 | spl48_142 | spl48_141 | spl48_140 | spl48_139 | spl48_138 | spl48_137 | spl48_136 | spl48_135 | spl48_134 | spl48_133 | spl48_132 | spl48_131 | spl48_130 | spl48_129 | ~ spl48_4), inference(avatar_split_clause, [], [f698, f727, f1323, f1333, f1343, f1353, f1363, f1373, f1383, f1393, f1403, f1413, f1423, f1433, f1443, f1453, f1463])).
fof(f698, plain, (~ (e13 = op1(e13, e13)) | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3), inference(duplicate_literal_removal, [], [f467])).
fof(f467, plain, (~ (e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13)) | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3), inference(cnf_transformation, [], [f39])).
fof(f1530, plain, (~ spl48_154 | ~ spl48_1 | spl48_16), inference(avatar_split_clause, [], [f453, f778, f715, f1527])).
fof(f453, plain, ((e13 = op1(e13, e10)) | ~ (e10 = op1(e13, e13)) | ~ sP0), inference(cnf_transformation, [], [f89])).
fof(f89, plain, ((((e13 = op1(e13, e10)) | ~ (e10 = op1(e13, e13))) & ((e12 = op1(e12, e10)) | ~ (e10 = op1(e12, e12))) & ((e11 = op1(e11, e10)) | ~ (e10 = op1(e11, e11))) & ((e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10))) & (e13 = op1(e10, op1(e10, e13))) & (e12 = op1(e10, op1(e10, e12))) & (e11 = op1(e10, op1(e10, e11))) & (e10 = op1(e10, op1(e10, e10)))) | ~ sP0), inference(nnf_transformation, [], [f21])).
fof(f1504, plain, (~ spl48_149 | ~ spl48_22 | spl48_27), inference(avatar_split_clause, [], [f444, f825, f804, f1500])).
fof(f444, plain, ((e12 = op1(e12, e11)) | ~ (e11 = op1(e12, e12)) | ~ sP1), inference(cnf_transformation, [], [f88])).
fof(f88, plain, ((((e13 = op1(e13, e11)) | ~ (e11 = op1(e13, e13))) & ((e12 = op1(e12, e11)) | ~ (e11 = op1(e12, e12))) & ((e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & ((e10 = op1(e10, e11)) | ~ (op1(e10, e10) = e11)) & (e13 = op1(e11, op1(e11, e13))) & (e12 = op1(e11, op1(e11, e12))) & (e11 = op1(e11, op1(e11, e11))) & (e10 = op1(e11, op1(e11, e10)))) | ~ sP1), inference(nnf_transformation, [], [f22])).
fof(f1493, plain, (~ spl48_144 | spl48_147), inference(avatar_split_clause, [], [f431, f1490, f1473])).
fof(f431, plain, ((e11 = op1(e12, op1(e12, e11))) | ~ sP2), inference(cnf_transformation, [], [f87])).
fof(f87, plain, ((((e13 = op1(e13, e12)) | ~ (e12 = op1(e13, e13))) & ((e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & ((e11 = op1(e11, e12)) | ~ (e12 = op1(e11, e11))) & ((e10 = op1(e10, e12)) | ~ (op1(e10, e10) = e12)) & (e13 = op1(e12, op1(e12, e13))) & (e12 = op1(e12, op1(e12, e12))) & (e11 = op1(e12, op1(e12, e11))) & (e10 = op1(e12, op1(e12, e10)))) | ~ sP2), inference(nnf_transformation, [], [f23])).
fof(f1471, plain, (~ spl48_143 | spl48_61), inference(avatar_split_clause, [], [f424, f970, f1463])).
fof(f970, plain, (spl48_61 <=> (e10 = op1(e10, e10))), introduced(avatar_definition, [new_symbols(naming, [spl48_61])])).
fof(f424, plain, ((e10 = op1(e10, e10)) | ~ sP3), inference(cnf_transformation, [], [f86])).
fof(f86, plain, (((~ (op1(e10, e10) = e13) | ~ (e10 = op1(e10, e13))) & (~ (op1(e10, e10) = e12) | ~ (e10 = op1(e10, e12))) & (~ (op1(e10, e10) = e11) | ~ (e10 = op1(e10, e11))) & (~ (e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10))) & (e10 = op1(e10, e10)) & (e10 = op1(e10, e10))) | ~ sP3), inference(nnf_transformation, [], [f24])).
fof(f1469, plain, (~ spl48_143 | ~ spl48_61), inference(avatar_split_clause, [], [f699, f970, f1463])).
fof(f699, plain, (~ (e10 = op1(e10, e10)) | ~ sP3), inference(duplicate_literal_removal, [], [f426])).
fof(f426, plain, (~ (e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10)) | ~ sP3), inference(cnf_transformation, [], [f86])).
fof(f1466, plain, (~ spl48_143 | ~ spl48_49 | ~ spl48_64), inference(avatar_split_clause, [], [f429, f982, f919, f1463])).
fof(f429, plain, (~ (op1(e10, e10) = e13) | ~ (e10 = op1(e10, e13)) | ~ sP3), inference(cnf_transformation, [], [f86])).
fof(f1461, plain, (~ spl48_142 | spl48_57), inference(avatar_split_clause, [], [f418, f953, f1453])).
fof(f418, plain, ((e10 = op1(e10, e11)) | ~ sP4), inference(cnf_transformation, [], [f85])).
fof(f85, plain, (((~ (e13 = op1(e11, e11)) | ~ (e11 = op1(e11, e13))) & (~ (e12 = op1(e11, e11)) | ~ (e11 = op1(e11, e12))) & (~ (e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & (~ (e10 = op1(e11, e11)) | ~ (e11 = op1(e11, e10))) & (e10 = op1(e11, e10)) & (e10 = op1(e10, e11))) | ~ sP4), inference(nnf_transformation, [], [f25])).
fof(f1451, plain, (~ spl48_141 | spl48_53), inference(avatar_split_clause, [], [f412, f936, f1443])).
fof(f412, plain, ((e10 = op1(e10, e12)) | ~ sP5), inference(cnf_transformation, [], [f84])).
fof(f84, plain, (((~ (e13 = op1(e12, e12)) | ~ (e12 = op1(e12, e13))) & (~ (e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & (~ (e11 = op1(e12, e12)) | ~ (e12 = op1(e12, e11))) & (~ (e10 = op1(e12, e12)) | ~ (e12 = op1(e12, e10))) & (e10 = op1(e12, e10)) & (e10 = op1(e10, e12))) | ~ sP5), inference(nnf_transformation, [], [f26])).
fof(f1450, plain, (~ spl48_141 | spl48_29), inference(avatar_split_clause, [], [f413, f834, f1443])).
fof(f413, plain, ((e10 = op1(e12, e10)) | ~ sP5), inference(cnf_transformation, [], [f84])).
fof(f1447, plain, (~ spl48_141 | ~ spl48_23), inference(avatar_split_clause, [], [f701, f808, f1443])).
fof(f701, plain, (~ (e12 = op1(e12, e12)) | ~ sP5), inference(duplicate_literal_removal, [], [f416])).
fof(f416, plain, (~ (e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12)) | ~ sP5), inference(cnf_transformation, [], [f84])).
fof(f1446, plain, (~ spl48_141 | ~ spl48_19 | ~ spl48_24), inference(avatar_split_clause, [], [f417, f812, f791, f1443])).
fof(f791, plain, (spl48_19 <=> (e12 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl48_19])])).
fof(f417, plain, (~ (e13 = op1(e12, e12)) | ~ (e12 = op1(e12, e13)) | ~ sP5), inference(cnf_transformation, [], [f84])).
fof(f1440, plain, (~ spl48_140 | spl48_13), inference(avatar_split_clause, [], [f407, f766, f1433])).
fof(f407, plain, ((e10 = op1(e13, e10)) | ~ sP6), inference(cnf_transformation, [], [f83])).
fof(f83, plain, (((~ (e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13))) & (~ (e12 = op1(e13, e13)) | ~ (e13 = op1(e13, e12))) & (~ (e11 = op1(e13, e13)) | ~ (e13 = op1(e13, e11))) & (~ (e10 = op1(e13, e13)) | ~ (e13 = op1(e13, e10))) & (e10 = op1(e13, e10)) & (e10 = op1(e10, e13))) | ~ sP6), inference(nnf_transformation, [], [f27])).
fof(f1430, plain, (~ spl48_139 | spl48_58), inference(avatar_split_clause, [], [f401, f957, f1423])).
fof(f401, plain, ((e11 = op1(e10, e11)) | ~ sP7), inference(cnf_transformation, [], [f82])).
fof(f82, plain, (((~ (op1(e10, e10) = e13) | ~ (e10 = op1(e10, e13))) & (~ (op1(e10, e10) = e12) | ~ (e10 = op1(e10, e12))) & (~ (op1(e10, e10) = e11) | ~ (e10 = op1(e10, e11))) & (~ (e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10))) & (e11 = op1(e10, e11)) & (e11 = op1(e11, e10))) | ~ sP7), inference(nnf_transformation, [], [f28])).
fof(f1429, plain, (~ spl48_139 | ~ spl48_61), inference(avatar_split_clause, [], [f703, f970, f1423])).
fof(f703, plain, (~ (e10 = op1(e10, e10)) | ~ sP7), inference(duplicate_literal_removal, [], [f402])).
fof(f402, plain, (~ (e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10)) | ~ sP7), inference(cnf_transformation, [], [f82])).
fof(f1426, plain, (~ spl48_139 | ~ spl48_49 | ~ spl48_64), inference(avatar_split_clause, [], [f405, f982, f919, f1423])).
fof(f405, plain, (~ (op1(e10, e10) = e13) | ~ (e10 = op1(e10, e13)) | ~ sP7), inference(cnf_transformation, [], [f82])).
fof(f1419, plain, (~ spl48_138 | ~ spl48_46 | ~ spl48_41), inference(avatar_split_clause, [], [f396, f885, f906, f1413])).
fof(f396, plain, (~ (e10 = op1(e11, e11)) | ~ (e11 = op1(e11, e10)) | ~ sP8), inference(cnf_transformation, [], [f81])).
fof(f81, plain, (((~ (e13 = op1(e11, e11)) | ~ (e11 = op1(e11, e13))) & (~ (e12 = op1(e11, e11)) | ~ (e11 = op1(e11, e12))) & (~ (e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & (~ (e10 = op1(e11, e11)) | ~ (e11 = op1(e11, e10))) & (e11 = op1(e11, e11)) & (e11 = op1(e11, e11))) | ~ sP8), inference(nnf_transformation, [], [f29])).
fof(f1411, plain, (~ spl48_137 | spl48_38), inference(avatar_split_clause, [], [f388, f872, f1403])).
fof(f388, plain, ((e11 = op1(e11, e12)) | ~ sP9), inference(cnf_transformation, [], [f80])).
fof(f80, plain, (((~ (e13 = op1(e12, e12)) | ~ (e12 = op1(e12, e13))) & (~ (e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & (~ (e11 = op1(e12, e12)) | ~ (e12 = op1(e12, e11))) & (~ (e10 = op1(e12, e12)) | ~ (e12 = op1(e12, e10))) & (e11 = op1(e12, e11)) & (e11 = op1(e11, e12))) | ~ sP9), inference(nnf_transformation, [], [f30])).
fof(f1401, plain, (~ spl48_136 | spl48_34), inference(avatar_split_clause, [], [f382, f855, f1393])).
fof(f382, plain, ((e11 = op1(e11, e13)) | ~ sP10), inference(cnf_transformation, [], [f79])).
fof(f79, plain, (((~ (e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13))) & (~ (e12 = op1(e13, e13)) | ~ (e13 = op1(e13, e12))) & (~ (e11 = op1(e13, e13)) | ~ (e13 = op1(e13, e11))) & (~ (e10 = op1(e13, e13)) | ~ (e13 = op1(e13, e10))) & (e11 = op1(e13, e11)) & (e11 = op1(e11, e13))) | ~ sP10), inference(nnf_transformation, [], [f31])).
fof(f1390, plain, (~ spl48_135 | spl48_55), inference(avatar_split_clause, [], [f377, f944, f1383])).
fof(f377, plain, ((e12 = op1(e10, e12)) | ~ sP11), inference(cnf_transformation, [], [f78])).
fof(f78, plain, (((~ (op1(e10, e10) = e13) | ~ (e10 = op1(e10, e13))) & (~ (op1(e10, e10) = e12) | ~ (e10 = op1(e10, e12))) & (~ (op1(e10, e10) = e11) | ~ (e10 = op1(e10, e11))) & (~ (e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10))) & (e12 = op1(e10, e12)) & (e12 = op1(e12, e10))) | ~ sP11), inference(nnf_transformation, [], [f32])).
fof(f1389, plain, (~ spl48_135 | ~ spl48_61), inference(avatar_split_clause, [], [f707, f970, f1383])).
fof(f707, plain, (~ (e10 = op1(e10, e10)) | ~ sP11), inference(duplicate_literal_removal, [], [f378])).
fof(f378, plain, (~ (e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10)) | ~ sP11), inference(cnf_transformation, [], [f78])).
fof(f1386, plain, (~ spl48_135 | ~ spl48_49 | ~ spl48_64), inference(avatar_split_clause, [], [f381, f982, f919, f1383])).
fof(f381, plain, (~ (op1(e10, e10) = e13) | ~ (e10 = op1(e10, e13)) | ~ sP11), inference(cnf_transformation, [], [f78])).
fof(f1379, plain, (~ spl48_134 | ~ spl48_46 | ~ spl48_41), inference(avatar_split_clause, [], [f372, f885, f906, f1373])).
fof(f372, plain, (~ (e10 = op1(e11, e11)) | ~ (e11 = op1(e11, e10)) | ~ sP12), inference(cnf_transformation, [], [f77])).
fof(f77, plain, (((~ (e13 = op1(e11, e11)) | ~ (e11 = op1(e11, e13))) & (~ (e12 = op1(e11, e11)) | ~ (e11 = op1(e11, e12))) & (~ (e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & (~ (e10 = op1(e11, e11)) | ~ (e11 = op1(e11, e10))) & (e12 = op1(e11, e12)) & (e12 = op1(e12, e11))) | ~ sP12), inference(nnf_transformation, [], [f33])).
fof(f1371, plain, (~ spl48_133 | spl48_23), inference(avatar_split_clause, [], [f364, f808, f1363])).
fof(f364, plain, ((e12 = op1(e12, e12)) | ~ sP13), inference(cnf_transformation, [], [f76])).
fof(f76, plain, (((~ (e13 = op1(e12, e12)) | ~ (e12 = op1(e12, e13))) & (~ (e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & (~ (e11 = op1(e12, e12)) | ~ (e12 = op1(e12, e11))) & (~ (e10 = op1(e12, e12)) | ~ (e12 = op1(e12, e10))) & (e12 = op1(e12, e12)) & (e12 = op1(e12, e12))) | ~ sP13), inference(nnf_transformation, [], [f34])).
fof(f1367, plain, (~ spl48_133 | ~ spl48_23), inference(avatar_split_clause, [], [f709, f808, f1363])).
fof(f709, plain, (~ (e12 = op1(e12, e12)) | ~ sP13), inference(duplicate_literal_removal, [], [f368])).
fof(f368, plain, (~ (e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12)) | ~ sP13), inference(cnf_transformation, [], [f76])).
fof(f1360, plain, (~ spl48_132 | spl48_7), inference(avatar_split_clause, [], [f359, f740, f1353])).
fof(f359, plain, ((e12 = op1(e13, e12)) | ~ sP14), inference(cnf_transformation, [], [f75])).
fof(f75, plain, (((~ (e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13))) & (~ (e12 = op1(e13, e13)) | ~ (e13 = op1(e13, e12))) & (~ (e11 = op1(e13, e13)) | ~ (e13 = op1(e13, e11))) & (~ (e10 = op1(e13, e13)) | ~ (e13 = op1(e13, e10))) & (e12 = op1(e13, e12)) & (e12 = op1(e12, e13))) | ~ sP14), inference(nnf_transformation, [], [f35])).
fof(f1351, plain, (~ spl48_131 | spl48_16), inference(avatar_split_clause, [], [f352, f778, f1343])).
fof(f352, plain, ((e13 = op1(e13, e10)) | ~ sP15), inference(cnf_transformation, [], [f74])).
fof(f74, plain, (((~ (op1(e10, e10) = e13) | ~ (e10 = op1(e10, e13))) & (~ (op1(e10, e10) = e12) | ~ (e10 = op1(e10, e12))) & (~ (op1(e10, e10) = e11) | ~ (e10 = op1(e10, e11))) & (~ (e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10))) & (e13 = op1(e10, e13)) & (e13 = op1(e13, e10))) | ~ sP15), inference(nnf_transformation, [], [f36])).
fof(f1339, plain, (~ spl48_130 | ~ spl48_46 | ~ spl48_41), inference(avatar_split_clause, [], [f348, f885, f906, f1333])).
fof(f348, plain, (~ (e10 = op1(e11, e11)) | ~ (e11 = op1(e11, e10)) | ~ sP16), inference(cnf_transformation, [], [f73])).
fof(f73, plain, (((~ (e13 = op1(e11, e11)) | ~ (e11 = op1(e11, e13))) & (~ (e12 = op1(e11, e11)) | ~ (e11 = op1(e11, e12))) & (~ (e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & (~ (e10 = op1(e11, e11)) | ~ (e11 = op1(e11, e10))) & (e13 = op1(e11, e13)) & (e13 = op1(e13, e11))) | ~ sP16), inference(nnf_transformation, [], [f37])).
fof(f1331, plain, (~ spl48_129 | spl48_8), inference(avatar_split_clause, [], [f340, f744, f1323])).
fof(f340, plain, ((e13 = op1(e13, e12)) | ~ sP17), inference(cnf_transformation, [], [f72])).
fof(f72, plain, (((~ (e13 = op1(e12, e12)) | ~ (e12 = op1(e12, e13))) & (~ (e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & (~ (e11 = op1(e12, e12)) | ~ (e12 = op1(e12, e11))) & (~ (e10 = op1(e12, e12)) | ~ (e12 = op1(e12, e10))) & (e13 = op1(e12, e13)) & (e13 = op1(e13, e12))) | ~ sP17), inference(nnf_transformation, [], [f38])).
fof(f1330, plain, (~ spl48_129 | spl48_20), inference(avatar_split_clause, [], [f341, f795, f1323])).
fof(f341, plain, ((e13 = op1(e12, e13)) | ~ sP17), inference(cnf_transformation, [], [f72])).
fof(f1321, plain, (spl48_125 | spl48_121 | spl48_117 | spl48_113), inference(avatar_split_clause, [], [f184, f1223, f1240, f1257, f1274])).
fof(f184, plain, ((e20 = op2(e20, e23)) | (e20 = op2(e20, e22)) | (e20 = op2(e20, e21)) | (e20 = op2(e20, e20))), inference(cnf_transformation, [], [f4])).
fof(f4, plain, (((e23 = op2(e23, e23)) | (e23 = op2(e22, e23)) | (e23 = op2(e21, e23)) | (e23 = op2(e20, e23))) & ((e23 = op2(e23, e23)) | (e23 = op2(e23, e22)) | (e23 = op2(e23, e21)) | (e23 = op2(e23, e20))) & ((e22 = op2(e23, e23)) | (e22 = op2(e22, e23)) | (e22 = op2(e21, e23)) | (e22 = op2(e20, e23))) & ((e22 = op2(e23, e23)) | (e22 = op2(e23, e22)) | (e22 = op2(e23, e21)) | (e22 = op2(e23, e20))) & ((e21 = op2(e23, e23)) | (e21 = op2(e22, e23)) | (e21 = op2(e21, e23)) | (e21 = op2(e20, e23))) & ((e21 = op2(e23, e23)) | (e21 = op2(e23, e22)) | (e21 = op2(e23, e21)) | (e21 = op2(e23, e20))) & ((e20 = op2(e23, e23)) | (e20 = op2(e22, e23)) | (e20 = op2(e21, e23)) | (e20 = op2(e20, e23))) & ((e20 = op2(e23, e23)) | (e20 = op2(e23, e22)) | (e20 = op2(e23, e21)) | (e20 = op2(e23, e20))) & ((e23 = op2(e23, e22)) | (e23 = op2(e22, e22)) | (e23 = op2(e21, e22)) | (e23 = op2(e20, e22))) & ((e23 = op2(e22, e23)) | (e23 = op2(e22, e22)) | (e23 = op2(e22, e21)) | (e23 = op2(e22, e20))) & ((e22 = op2(e23, e22)) | (e22 = op2(e22, e22)) | (e22 = op2(e21, e22)) | (e22 = op2(e20, e22))) & ((e22 = op2(e22, e23)) | (e22 = op2(e22, e22)) | (e22 = op2(e22, e21)) | (e22 = op2(e22, e20))) & ((e21 = op2(e23, e22)) | (e21 = op2(e22, e22)) | (e21 = op2(e21, e22)) | (e21 = op2(e20, e22))) & ((e21 = op2(e22, e23)) | (e21 = op2(e22, e22)) | (e21 = op2(e22, e21)) | (e21 = op2(e22, e20))) & ((e20 = op2(e23, e22)) | (e20 = op2(e22, e22)) | (e20 = op2(e21, e22)) | (e20 = op2(e20, e22))) & ((e20 = op2(e22, e23)) | (e20 = op2(e22, e22)) | (e20 = op2(e22, e21)) | (e20 = op2(e22, e20))) & ((e23 = op2(e23, e21)) | (e23 = op2(e22, e21)) | (e23 = op2(e21, e21)) | (e23 = op2(e20, e21))) & ((e23 = op2(e21, e23)) | (e23 = op2(e21, e22)) | (e23 = op2(e21, e21)) | (e23 = op2(e21, e20))) & ((e22 = op2(e23, e21)) | (e22 = op2(e22, e21)) | (e22 = op2(e21, e21)) | (e22 = op2(e20, e21))) & ((e22 = op2(e21, e23)) | (e22 = op2(e21, e22)) | (e22 = op2(e21, e21)) | (e22 = op2(e21, e20))) & ((e21 = op2(e23, e21)) | (e21 = op2(e22, e21)) | (e21 = op2(e21, e21)) | (e21 = op2(e20, e21))) & ((e21 = op2(e21, e23)) | (e21 = op2(e21, e22)) | (e21 = op2(e21, e21)) | (e21 = op2(e21, e20))) & ((e20 = op2(e23, e21)) | (e20 = op2(e22, e21)) | (e20 = op2(e21, e21)) | (e20 = op2(e20, e21))) & ((e20 = op2(e21, e23)) | (e20 = op2(e21, e22)) | (e20 = op2(e21, e21)) | (e20 = op2(e21, e20))) & ((e23 = op2(e23, e20)) | (e23 = op2(e22, e20)) | (e23 = op2(e21, e20)) | (op2(e20, e20) = e23)) & ((e23 = op2(e20, e23)) | (e23 = op2(e20, e22)) | (e23 = op2(e20, e21)) | (op2(e20, e20) = e23)) & ((e22 = op2(e23, e20)) | (e22 = op2(e22, e20)) | (e22 = op2(e21, e20)) | (op2(e20, e20) = e22)) & ((e22 = op2(e20, e23)) | (e22 = op2(e20, e22)) | (e22 = op2(e20, e21)) | (op2(e20, e20) = e22)) & ((e21 = op2(e23, e20)) | (e21 = op2(e22, e20)) | (e21 = op2(e21, e20)) | (op2(e20, e20) = e21)) & ((e21 = op2(e20, e23)) | (e21 = op2(e20, e22)) | (e21 = op2(e20, e21)) | (op2(e20, e20) = e21)) & ((e20 = op2(e23, e20)) | (e20 = op2(e22, e20)) | (e20 = op2(e21, e20)) | (e20 = op2(e20, e20))) & ((e20 = op2(e20, e23)) | (e20 = op2(e20, e22)) | (e20 = op2(e20, e21)) | (e20 = op2(e20, e20)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG126+1.p', ax4)).
fof(f1308, plain, (spl48_123 | spl48_107 | spl48_91 | spl48_75), inference(avatar_split_clause, [], [f197, f1061, f1129, f1197, f1265])).
fof(f197, plain, ((e22 = op2(e23, e21)) | (e22 = op2(e22, e21)) | (e22 = op2(e21, e21)) | (e22 = op2(e20, e21))), inference(cnf_transformation, [], [f4])).
fof(f1306, plain, (spl48_124 | spl48_108 | spl48_92 | spl48_76), inference(avatar_split_clause, [], [f199, f1065, f1133, f1201, f1269])).
fof(f199, plain, ((e23 = op2(e23, e21)) | (e23 = op2(e22, e21)) | (e23 = op2(e21, e21)) | (e23 = op2(e20, e21))), inference(cnf_transformation, [], [f4])).
fof(f1301, plain, (spl48_95 | spl48_91 | spl48_87 | spl48_83), inference(avatar_split_clause, [], [f204, f1095, f1112, f1129, f1146])).
fof(f204, plain, ((e22 = op2(e22, e23)) | (e22 = op2(e22, e22)) | (e22 = op2(e22, e21)) | (e22 = op2(e22, e20))), inference(cnf_transformation, [], [f4])).
fof(f1300, plain, (spl48_119 | spl48_103 | spl48_87 | spl48_71), inference(avatar_split_clause, [], [f205, f1044, f1112, f1180, f1248])).
fof(f205, plain, ((e22 = op2(e23, e22)) | (e22 = op2(e22, e22)) | (e22 = op2(e21, e22)) | (e22 = op2(e20, e22))), inference(cnf_transformation, [], [f4])).
fof(f1298, plain, (spl48_120 | spl48_104 | spl48_88 | spl48_72), inference(avatar_split_clause, [], [f207, f1048, f1116, f1184, f1252])).
fof(f207, plain, ((e23 = op2(e23, e22)) | (e23 = op2(e22, e22)) | (e23 = op2(e21, e22)) | (e23 = op2(e20, e22))), inference(cnf_transformation, [], [f4])).
fof(f1296, plain, (spl48_113 | spl48_97 | spl48_81 | spl48_65), inference(avatar_split_clause, [], [f209, f1019, f1087, f1155, f1223])).
fof(f209, plain, ((e20 = op2(e23, e23)) | (e20 = op2(e22, e23)) | (e20 = op2(e21, e23)) | (e20 = op2(e20, e23))), inference(cnf_transformation, [], [f4])).
fof(f1295, plain, (spl48_78 | spl48_74 | spl48_70 | spl48_66), inference(avatar_split_clause, [], [f210, f1023, f1040, f1057, f1074])).
fof(f210, plain, ((e21 = op2(e23, e23)) | (e21 = op2(e23, e22)) | (e21 = op2(e23, e21)) | (e21 = op2(e23, e20))), inference(cnf_transformation, [], [f4])).
fof(f1294, plain, (spl48_114 | spl48_98 | spl48_82 | spl48_66), inference(avatar_split_clause, [], [f211, f1023, f1091, f1159, f1227])).
fof(f211, plain, ((e21 = op2(e23, e23)) | (e21 = op2(e22, e23)) | (e21 = op2(e21, e23)) | (e21 = op2(e20, e23))), inference(cnf_transformation, [], [f4])).
fof(f1272, plain, (spl48_121 | spl48_122 | spl48_123 | spl48_124), inference(avatar_split_clause, [], [f169, f1269, f1265, f1261, f1257])).
fof(f169, plain, ((e23 = op2(e20, e21)) | (e22 = op2(e20, e21)) | (e21 = op2(e20, e21)) | (e20 = op2(e20, e21))), inference(cnf_transformation, [], [f3])).
fof(f3, plain, (((e23 = op2(e23, e23)) | (e22 = op2(e23, e23)) | (e21 = op2(e23, e23)) | (e20 = op2(e23, e23))) & ((e23 = op2(e23, e22)) | (e22 = op2(e23, e22)) | (e21 = op2(e23, e22)) | (e20 = op2(e23, e22))) & ((e23 = op2(e23, e21)) | (e22 = op2(e23, e21)) | (e21 = op2(e23, e21)) | (e20 = op2(e23, e21))) & ((e23 = op2(e23, e20)) | (e22 = op2(e23, e20)) | (e21 = op2(e23, e20)) | (e20 = op2(e23, e20))) & ((e23 = op2(e22, e23)) | (e22 = op2(e22, e23)) | (e21 = op2(e22, e23)) | (e20 = op2(e22, e23))) & ((e23 = op2(e22, e22)) | (e22 = op2(e22, e22)) | (e21 = op2(e22, e22)) | (e20 = op2(e22, e22))) & ((e23 = op2(e22, e21)) | (e22 = op2(e22, e21)) | (e21 = op2(e22, e21)) | (e20 = op2(e22, e21))) & ((e23 = op2(e22, e20)) | (e22 = op2(e22, e20)) | (e21 = op2(e22, e20)) | (e20 = op2(e22, e20))) & ((e23 = op2(e21, e23)) | (e22 = op2(e21, e23)) | (e21 = op2(e21, e23)) | (e20 = op2(e21, e23))) & ((e23 = op2(e21, e22)) | (e22 = op2(e21, e22)) | (e21 = op2(e21, e22)) | (e20 = op2(e21, e22))) & ((e23 = op2(e21, e21)) | (e22 = op2(e21, e21)) | (e21 = op2(e21, e21)) | (e20 = op2(e21, e21))) & ((e23 = op2(e21, e20)) | (e22 = op2(e21, e20)) | (e21 = op2(e21, e20)) | (e20 = op2(e21, e20))) & ((e23 = op2(e20, e23)) | (e22 = op2(e20, e23)) | (e21 = op2(e20, e23)) | (e20 = op2(e20, e23))) & ((e23 = op2(e20, e22)) | (e22 = op2(e20, e22)) | (e21 = op2(e20, e22)) | (e20 = op2(e20, e22))) & ((e23 = op2(e20, e21)) | (e22 = op2(e20, e21)) | (e21 = op2(e20, e21)) | (e20 = op2(e20, e21))) & ((op2(e20, e20) = e23) | (op2(e20, e20) = e22) | (op2(e20, e20) = e21) | (e20 = op2(e20, e20)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG126+1.p', ax3)).
fof(f1238, plain, (spl48_113 | spl48_114 | spl48_115 | spl48_116), inference(avatar_split_clause, [], [f171, f1235, f1231, f1227, f1223])).
fof(f171, plain, ((e23 = op2(e20, e23)) | (e22 = op2(e20, e23)) | (e21 = op2(e20, e23)) | (e20 = op2(e20, e23))), inference(cnf_transformation, [], [f3])).
fof(f1221, plain, (spl48_109 | spl48_110 | spl48_111 | spl48_112), inference(avatar_split_clause, [], [f172, f1218, f1214, f1210, f1206])).
fof(f172, plain, ((e23 = op2(e21, e20)) | (e22 = op2(e21, e20)) | (e21 = op2(e21, e20)) | (e20 = op2(e21, e20))), inference(cnf_transformation, [], [f3])).
fof(f1170, plain, (spl48_97 | spl48_98 | spl48_99 | spl48_100), inference(avatar_split_clause, [], [f175, f1167, f1163, f1159, f1155])).
fof(f175, plain, ((e23 = op2(e21, e23)) | (e22 = op2(e21, e23)) | (e21 = op2(e21, e23)) | (e20 = op2(e21, e23))), inference(cnf_transformation, [], [f3])).
fof(f1153, plain, (spl48_93 | spl48_94 | spl48_95 | spl48_96), inference(avatar_split_clause, [], [f176, f1150, f1146, f1142, f1138])).
fof(f176, plain, ((e23 = op2(e22, e20)) | (e22 = op2(e22, e20)) | (e21 = op2(e22, e20)) | (e20 = op2(e22, e20))), inference(cnf_transformation, [], [f3])).
fof(f1119, plain, (spl48_85 | spl48_86 | spl48_87 | spl48_88), inference(avatar_split_clause, [], [f178, f1116, f1112, f1108, f1104])).
fof(f178, plain, ((e23 = op2(e22, e22)) | (e22 = op2(e22, e22)) | (e21 = op2(e22, e22)) | (e20 = op2(e22, e22))), inference(cnf_transformation, [], [f3])).
fof(f1102, plain, (spl48_81 | spl48_82 | spl48_83 | spl48_84), inference(avatar_split_clause, [], [f179, f1099, f1095, f1091, f1087])).
fof(f179, plain, ((e23 = op2(e22, e23)) | (e22 = op2(e22, e23)) | (e21 = op2(e22, e23)) | (e20 = op2(e22, e23))), inference(cnf_transformation, [], [f3])).
fof(f1051, plain, (spl48_69 | spl48_70 | spl48_71 | spl48_72), inference(avatar_split_clause, [], [f182, f1048, f1044, f1040, f1036])).
fof(f182, plain, ((e23 = op2(e23, e22)) | (e22 = op2(e23, e22)) | (e21 = op2(e23, e22)) | (e20 = op2(e23, e22))), inference(cnf_transformation, [], [f3])).
fof(f1017, plain, (spl48_61 | spl48_57 | spl48_53 | spl48_49), inference(avatar_split_clause, [], [f136, f919, f936, f953, f970])).
fof(f136, plain, ((e10 = op1(e10, e13)) | (e10 = op1(e10, e12)) | (e10 = op1(e10, e11)) | (e10 = op1(e10, e10))), inference(cnf_transformation, [], [f2])).
fof(f2, plain, (((e13 = op1(e13, e13)) | (e13 = op1(e12, e13)) | (e13 = op1(e11, e13)) | (e13 = op1(e10, e13))) & ((e13 = op1(e13, e13)) | (e13 = op1(e13, e12)) | (e13 = op1(e13, e11)) | (e13 = op1(e13, e10))) & ((e12 = op1(e13, e13)) | (e12 = op1(e12, e13)) | (e12 = op1(e11, e13)) | (e12 = op1(e10, e13))) & ((e12 = op1(e13, e13)) | (e12 = op1(e13, e12)) | (e12 = op1(e13, e11)) | (e12 = op1(e13, e10))) & ((e11 = op1(e13, e13)) | (e11 = op1(e12, e13)) | (e11 = op1(e11, e13)) | (e11 = op1(e10, e13))) & ((e11 = op1(e13, e13)) | (e11 = op1(e13, e12)) | (e11 = op1(e13, e11)) | (e11 = op1(e13, e10))) & ((e10 = op1(e13, e13)) | (e10 = op1(e12, e13)) | (e10 = op1(e11, e13)) | (e10 = op1(e10, e13))) & ((e10 = op1(e13, e13)) | (e10 = op1(e13, e12)) | (e10 = op1(e13, e11)) | (e10 = op1(e13, e10))) & ((e13 = op1(e13, e12)) | (e13 = op1(e12, e12)) | (e13 = op1(e11, e12)) | (e13 = op1(e10, e12))) & ((e13 = op1(e12, e13)) | (e13 = op1(e12, e12)) | (e13 = op1(e12, e11)) | (e13 = op1(e12, e10))) & ((e12 = op1(e13, e12)) | (e12 = op1(e12, e12)) | (e12 = op1(e11, e12)) | (e12 = op1(e10, e12))) & ((e12 = op1(e12, e13)) | (e12 = op1(e12, e12)) | (e12 = op1(e12, e11)) | (e12 = op1(e12, e10))) & ((e11 = op1(e13, e12)) | (e11 = op1(e12, e12)) | (e11 = op1(e11, e12)) | (e11 = op1(e10, e12))) & ((e11 = op1(e12, e13)) | (e11 = op1(e12, e12)) | (e11 = op1(e12, e11)) | (e11 = op1(e12, e10))) & ((e10 = op1(e13, e12)) | (e10 = op1(e12, e12)) | (e10 = op1(e11, e12)) | (e10 = op1(e10, e12))) & ((e10 = op1(e12, e13)) | (e10 = op1(e12, e12)) | (e10 = op1(e12, e11)) | (e10 = op1(e12, e10))) & ((e13 = op1(e13, e11)) | (e13 = op1(e12, e11)) | (e13 = op1(e11, e11)) | (e13 = op1(e10, e11))) & ((e13 = op1(e11, e13)) | (e13 = op1(e11, e12)) | (e13 = op1(e11, e11)) | (e13 = op1(e11, e10))) & ((e12 = op1(e13, e11)) | (e12 = op1(e12, e11)) | (e12 = op1(e11, e11)) | (e12 = op1(e10, e11))) & ((e12 = op1(e11, e13)) | (e12 = op1(e11, e12)) | (e12 = op1(e11, e11)) | (e12 = op1(e11, e10))) & ((e11 = op1(e13, e11)) | (e11 = op1(e12, e11)) | (e11 = op1(e11, e11)) | (e11 = op1(e10, e11))) & ((e11 = op1(e11, e13)) | (e11 = op1(e11, e12)) | (e11 = op1(e11, e11)) | (e11 = op1(e11, e10))) & ((e10 = op1(e13, e11)) | (e10 = op1(e12, e11)) | (e10 = op1(e11, e11)) | (e10 = op1(e10, e11))) & ((e10 = op1(e11, e13)) | (e10 = op1(e11, e12)) | (e10 = op1(e11, e11)) | (e10 = op1(e11, e10))) & ((e13 = op1(e13, e10)) | (e13 = op1(e12, e10)) | (e13 = op1(e11, e10)) | (op1(e10, e10) = e13)) & ((e13 = op1(e10, e13)) | (e13 = op1(e10, e12)) | (e13 = op1(e10, e11)) | (op1(e10, e10) = e13)) & ((e12 = op1(e13, e10)) | (e12 = op1(e12, e10)) | (e12 = op1(e11, e10)) | (op1(e10, e10) = e12)) & ((e12 = op1(e10, e13)) | (e12 = op1(e10, e12)) | (e12 = op1(e10, e11)) | (op1(e10, e10) = e12)) & ((e11 = op1(e13, e10)) | (e11 = op1(e12, e10)) | (e11 = op1(e11, e10)) | (op1(e10, e10) = e11)) & ((e11 = op1(e10, e13)) | (e11 = op1(e10, e12)) | (e11 = op1(e10, e11)) | (op1(e10, e10) = e11)) & ((e10 = op1(e13, e10)) | (e10 = op1(e12, e10)) | (e10 = op1(e11, e10)) | (e10 = op1(e10, e10))) & ((e10 = op1(e10, e13)) | (e10 = op1(e10, e12)) | (e10 = op1(e10, e11)) | (e10 = op1(e10, e10)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG126+1.p', ax2)).
fof(f1016, plain, (spl48_61 | spl48_45 | spl48_29 | spl48_13), inference(avatar_split_clause, [], [f137, f766, f834, f902, f970])).
fof(f137, plain, ((e10 = op1(e13, e10)) | (e10 = op1(e12, e10)) | (e10 = op1(e11, e10)) | (e10 = op1(e10, e10))), inference(cnf_transformation, [], [f2])).
fof(f1005, plain, (spl48_47 | spl48_43 | spl48_39 | spl48_35), inference(avatar_split_clause, [], [f148, f859, f876, f893, f910])).
fof(f148, plain, ((e12 = op1(e11, e13)) | (e12 = op1(e11, e12)) | (e12 = op1(e11, e11)) | (e12 = op1(e11, e10))), inference(cnf_transformation, [], [f2])).
fof(f1004, plain, (spl48_59 | spl48_43 | spl48_27 | spl48_11), inference(avatar_split_clause, [], [f149, f757, f825, f893, f961])).
fof(f149, plain, ((e12 = op1(e13, e11)) | (e12 = op1(e12, e11)) | (e12 = op1(e11, e11)) | (e12 = op1(e10, e11))), inference(cnf_transformation, [], [f2])).
fof(f1002, plain, (spl48_60 | spl48_44 | spl48_28 | spl48_12), inference(avatar_split_clause, [], [f151, f761, f829, f897, f965])).
fof(f151, plain, ((e13 = op1(e13, e11)) | (e13 = op1(e12, e11)) | (e13 = op1(e11, e11)) | (e13 = op1(e10, e11))), inference(cnf_transformation, [], [f2])).
fof(f999, plain, (spl48_30 | spl48_26 | spl48_22 | spl48_18), inference(avatar_split_clause, [], [f154, f787, f804, f821, f838])).
fof(f154, plain, ((e11 = op1(e12, e13)) | (e11 = op1(e12, e12)) | (e11 = op1(e12, e11)) | (e11 = op1(e12, e10))), inference(cnf_transformation, [], [f2])).
fof(f996, plain, (spl48_55 | spl48_39 | spl48_23 | spl48_7), inference(avatar_split_clause, [], [f157, f740, f808, f876, f944])).
fof(f157, plain, ((e12 = op1(e13, e12)) | (e12 = op1(e12, e12)) | (e12 = op1(e11, e12)) | (e12 = op1(e10, e12))), inference(cnf_transformation, [], [f2])).
fof(f993, plain, (spl48_13 | spl48_9 | spl48_5 | spl48_1), inference(avatar_split_clause, [], [f160, f715, f732, f749, f766])).
fof(f160, plain, ((e10 = op1(e13, e13)) | (e10 = op1(e13, e12)) | (e10 = op1(e13, e11)) | (e10 = op1(e13, e10))), inference(cnf_transformation, [], [f2])).
fof(f988, plain, (spl48_51 | spl48_35 | spl48_19 | spl48_3), inference(avatar_split_clause, [], [f165, f723, f791, f859, f927])).
fof(f165, plain, ((e12 = op1(e13, e13)) | (e12 = op1(e12, e13)) | (e12 = op1(e11, e13)) | (e12 = op1(e10, e13))), inference(cnf_transformation, [], [f2])).
fof(f968, plain, (spl48_57 | spl48_58 | spl48_59 | spl48_60), inference(avatar_split_clause, [], [f121, f965, f961, f957, f953])).
fof(f121, plain, ((e13 = op1(e10, e11)) | (e12 = op1(e10, e11)) | (e11 = op1(e10, e11)) | (e10 = op1(e10, e11))), inference(cnf_transformation, [], [f1])).
fof(f1, plain, (((e13 = op1(e13, e13)) | (e12 = op1(e13, e13)) | (e11 = op1(e13, e13)) | (e10 = op1(e13, e13))) & ((e13 = op1(e13, e12)) | (e12 = op1(e13, e12)) | (e11 = op1(e13, e12)) | (e10 = op1(e13, e12))) & ((e13 = op1(e13, e11)) | (e12 = op1(e13, e11)) | (e11 = op1(e13, e11)) | (e10 = op1(e13, e11))) & ((e13 = op1(e13, e10)) | (e12 = op1(e13, e10)) | (e11 = op1(e13, e10)) | (e10 = op1(e13, e10))) & ((e13 = op1(e12, e13)) | (e12 = op1(e12, e13)) | (e11 = op1(e12, e13)) | (e10 = op1(e12, e13))) & ((e13 = op1(e12, e12)) | (e12 = op1(e12, e12)) | (e11 = op1(e12, e12)) | (e10 = op1(e12, e12))) & ((e13 = op1(e12, e11)) | (e12 = op1(e12, e11)) | (e11 = op1(e12, e11)) | (e10 = op1(e12, e11))) & ((e13 = op1(e12, e10)) | (e12 = op1(e12, e10)) | (e11 = op1(e12, e10)) | (e10 = op1(e12, e10))) & ((e13 = op1(e11, e13)) | (e12 = op1(e11, e13)) | (e11 = op1(e11, e13)) | (e10 = op1(e11, e13))) & ((e13 = op1(e11, e12)) | (e12 = op1(e11, e12)) | (e11 = op1(e11, e12)) | (e10 = op1(e11, e12))) & ((e13 = op1(e11, e11)) | (e12 = op1(e11, e11)) | (e11 = op1(e11, e11)) | (e10 = op1(e11, e11))) & ((e13 = op1(e11, e10)) | (e12 = op1(e11, e10)) | (e11 = op1(e11, e10)) | (e10 = op1(e11, e10))) & ((e13 = op1(e10, e13)) | (e12 = op1(e10, e13)) | (e11 = op1(e10, e13)) | (e10 = op1(e10, e13))) & ((e13 = op1(e10, e12)) | (e12 = op1(e10, e12)) | (e11 = op1(e10, e12)) | (e10 = op1(e10, e12))) & ((e13 = op1(e10, e11)) | (e12 = op1(e10, e11)) | (e11 = op1(e10, e11)) | (e10 = op1(e10, e11))) & ((op1(e10, e10) = e13) | (op1(e10, e10) = e12) | (op1(e10, e10) = e11) | (e10 = op1(e10, e10)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG126+1.p', ax1)).
fof(f934, plain, (spl48_49 | spl48_50 | spl48_51 | spl48_52), inference(avatar_split_clause, [], [f123, f931, f927, f923, f919])).
fof(f123, plain, ((e13 = op1(e10, e13)) | (e12 = op1(e10, e13)) | (e11 = op1(e10, e13)) | (e10 = op1(e10, e13))), inference(cnf_transformation, [], [f1])).
fof(f917, plain, (spl48_45 | spl48_46 | spl48_47 | spl48_48), inference(avatar_split_clause, [], [f124, f914, f910, f906, f902])).
fof(f124, plain, ((e13 = op1(e11, e10)) | (e12 = op1(e11, e10)) | (e11 = op1(e11, e10)) | (e10 = op1(e11, e10))), inference(cnf_transformation, [], [f1])).
fof(f866, plain, (spl48_33 | spl48_34 | spl48_35 | spl48_36), inference(avatar_split_clause, [], [f127, f863, f859, f855, f851])).
fof(f127, plain, ((e13 = op1(e11, e13)) | (e12 = op1(e11, e13)) | (e11 = op1(e11, e13)) | (e10 = op1(e11, e13))), inference(cnf_transformation, [], [f1])).
fof(f832, plain, (spl48_25 | spl48_26 | spl48_27 | spl48_28), inference(avatar_split_clause, [], [f129, f829, f825, f821, f817])).
fof(f129, plain, ((e13 = op1(e12, e11)) | (e12 = op1(e12, e11)) | (e11 = op1(e12, e11)) | (e10 = op1(e12, e11))), inference(cnf_transformation, [], [f1])).
fof(f815, plain, (spl48_21 | spl48_22 | spl48_23 | spl48_24), inference(avatar_split_clause, [], [f130, f812, f808, f804, f800])).
fof(f130, plain, ((e13 = op1(e12, e12)) | (e12 = op1(e12, e12)) | (e11 = op1(e12, e12)) | (e10 = op1(e12, e12))), inference(cnf_transformation, [], [f1])).
fof(f798, plain, (spl48_17 | spl48_18 | spl48_19 | spl48_20), inference(avatar_split_clause, [], [f131, f795, f791, f787, f783])).
fof(f131, plain, ((e13 = op1(e12, e13)) | (e12 = op1(e12, e13)) | (e11 = op1(e12, e13)) | (e10 = op1(e12, e13))), inference(cnf_transformation, [], [f1])).
fof(f764, plain, (spl48_9 | spl48_10 | spl48_11 | spl48_12), inference(avatar_split_clause, [], [f133, f761, f757, f753, f749])).
fof(f133, plain, ((e13 = op1(e13, e11)) | (e12 = op1(e13, e11)) | (e11 = op1(e13, e11)) | (e10 = op1(e13, e11))), inference(cnf_transformation, [], [f1])).
fof(f747, plain, (spl48_5 | spl48_6 | spl48_7 | spl48_8), inference(avatar_split_clause, [], [f134, f744, f740, f736, f732])).
fof(f134, plain, ((e13 = op1(e13, e12)) | (e12 = op1(e13, e12)) | (e11 = op1(e13, e12)) | (e10 = op1(e13, e12))), inference(cnf_transformation, [], [f1])).