fof(f3099, plain, $false, inference(avatar_sat_refutation, [], [f674, f708, f759, f761, f762, f763, f768, f771, f775, f780, f781, f783, f791, f1012, f1065, f1073, f1074, f1075, f1079, f1080, f1083, f1084, f1085, f1093, f1094, f1095, f1096, f1104, f1114, f1117, f1120, f1138, f1146, f1148, f1151, f1152, f1159, f1160, f1168, f1178, f1181, f1184, f1202, f1210, f1216, f1223, f1224, f1226, f1279, f1293, f1509, f1891, f1908, f1919, f1981, f2007, f2014, f2049, f2050, f2082, f2146, f2157, f2160, f2184, f2197, f2201, f2235, f2238, f2256, f2295, f2316, f2326, f2345, f2347, f2350, f2379, f2380, f2384, f2400, f2406, f2414, f2419, f2430, f2461, f2472, f2479, f2482, f2500, f2513, f2564, f2584, f2621, f2633, f2658, f2680, f2701, f2711, f2767, f2802, f2815, f2837, f2859, f2878, f2893, f2936, f2944, f2968, f2976, f2989, f3007, f3022, f3037, f3055, f3067, f3084, f3096])).
fof(f3096, plain, (~ spl12_62 | ~ spl12_130 | ~ spl12_155 | spl12_172), inference(avatar_contradiction_clause, [], [f3095])).
fof(f3095, plain, ($false | (~ spl12_62 | ~ spl12_130 | ~ spl12_155 | spl12_172)), inference(subsumption_resolution, [], [f3094, f1408])).
fof(f1408, plain, ((e21 = h1(e11)) | ~ spl12_155), inference(avatar_component_clause, [], [f1407])).
fof(f1407, plain, (spl12_155 <=> (e21 = h1(e11))), introduced(avatar_definition, [new_symbols(naming, [spl12_155])])).
fof(f3094, plain, (~ (e21 = h1(e11)) | (~ spl12_62 | ~ spl12_130 | spl12_172)), inference(forward_demodulation, [], [f3093, f1277])).
fof(f1277, plain, ((e21 = h4(e11)) | ~ spl12_130), inference(avatar_component_clause, [], [f1276])).
fof(f1276, plain, (spl12_130 <=> (e21 = h4(e11))), introduced(avatar_definition, [new_symbols(naming, [spl12_130])])).
fof(f3093, plain, (~ (h1(e11) = h4(e11)) | (~ spl12_62 | spl12_172)), inference(forward_demodulation, [], [f1508, f750])).
fof(f750, plain, ((op1(e10, e10) = e11) | ~ spl12_62), inference(avatar_component_clause, [], [f748])).
fof(f748, plain, (spl12_62 <=> (op1(e10, e10) = e11)), introduced(avatar_definition, [new_symbols(naming, [spl12_62])])).
fof(f1508, plain, (~ (h1(e11) = h4(op1(e10, e10))) | spl12_172), inference(avatar_component_clause, [], [f1506])).
fof(f1506, plain, (spl12_172 <=> (h1(e11) = h4(op1(e10, e10)))), introduced(avatar_definition, [new_symbols(naming, [spl12_172])])).
fof(f3084, plain, (~ spl12_53 | ~ spl12_117 | spl12_170), inference(avatar_contradiction_clause, [], [f3083])).
fof(f3083, plain, ($false | (~ spl12_53 | ~ spl12_117 | spl12_170)), inference(subsumption_resolution, [], [f3082, f1016])).
fof(f1016, plain, ((e20 = op2(e20, e22)) | ~ spl12_117), inference(avatar_component_clause, [], [f1014])).
fof(f1014, plain, (spl12_117 <=> (e20 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl12_117])])).
fof(f3082, plain, (~ (e20 = op2(e20, e22)) | (~ spl12_53 | spl12_170)), inference(forward_demodulation, [], [f3081, f1266])).
fof(f1266, plain, (e20 = h4(e10)), inference(forward_demodulation, [], [f1265, f1263])).
fof(f1263, plain, (e20 = op2(e23, h4(e11))), inference(backward_demodulation, [], [f397, f414])).
fof(f414, plain, (op2(e23, e23) = h4(e11)), inference(cnf_transformation, [], [f17])).
fof(f17, plain, ((op2(op2(e23, op2(e23, e23)), op2(e23, e23)) = h4(e12)) & (op2(e23, e23) = h4(e11)) & (op2(e23, op2(e23, e23)) = h4(e10)) & (e23 = h4(e13))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG111+1.p', ax17)).
fof(f397, plain, (e20 = op2(e23, op2(e23, e23))), inference(cnf_transformation, [], [f13])).
fof(f13, plain, ((e22 = op2(op2(e23, op2(e23, e23)), op2(e23, e23))) & (e21 = op2(e23, e23)) & (e20 = op2(e23, op2(e23, e23)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG111+1.p', ax13)).
fof(f1265, plain, (h4(e10) = op2(e23, h4(e11))), inference(forward_demodulation, [], [f413, f414])).
fof(f413, plain, (op2(e23, op2(e23, e23)) = h4(e10)), inference(cnf_transformation, [], [f17])).
fof(f3081, plain, (~ (op2(e20, e22) = h4(e10)) | (~ spl12_53 | spl12_170)), inference(forward_demodulation, [], [f1500, f712])).
fof(f712, plain, ((e10 = op1(e10, e12)) | ~ spl12_53), inference(avatar_component_clause, [], [f710])).
fof(f710, plain, (spl12_53 <=> (e10 = op1(e10, e12))), introduced(avatar_definition, [new_symbols(naming, [spl12_53])])).
fof(f1500, plain, (~ (op2(e20, e22) = h4(op1(e10, e12))) | spl12_170), inference(avatar_component_clause, [], [f1498])).
fof(f1498, plain, (spl12_170 <=> (op2(e20, e22) = h4(op1(e10, e12)))), introduced(avatar_definition, [new_symbols(naming, [spl12_170])])).
fof(f3067, plain, (~ spl12_59 | spl12_171), inference(avatar_contradiction_clause, [], [f3066])).
fof(f3066, plain, ($false | (~ spl12_59 | spl12_171)), inference(subsumption_resolution, [], [f3065, f1256])).
fof(f1256, plain, (e22 = h4(e12)), inference(forward_demodulation, [], [f1255, f1227])).
fof(f1227, plain, (e22 = op2(e20, op2(e23, e23))), inference(backward_demodulation, [], [f399, f397])).
fof(f399, plain, (e22 = op2(op2(e23, op2(e23, e23)), op2(e23, e23))), inference(cnf_transformation, [], [f13])).
fof(f1255, plain, (h4(e12) = op2(e20, op2(e23, e23))), inference(forward_demodulation, [], [f415, f397])).
fof(f415, plain, (op2(op2(e23, op2(e23, e23)), op2(e23, e23)) = h4(e12)), inference(cnf_transformation, [], [f17])).
fof(f3065, plain, (~ (e22 = h4(e12)) | (~ spl12_59 | spl12_171)), inference(forward_demodulation, [], [f1504, f737])).
fof(f737, plain, ((e12 = op1(e10, e11)) | ~ spl12_59), inference(avatar_component_clause, [], [f735])).
fof(f735, plain, (spl12_59 <=> (e12 = op1(e10, e11))), introduced(avatar_definition, [new_symbols(naming, [spl12_59])])).
fof(f1504, plain, (~ (e22 = h4(op1(e10, e11))) | spl12_171), inference(avatar_component_clause, [], [f1502])).
fof(f1502, plain, (spl12_171 <=> (e22 = h4(op1(e10, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl12_171])])).
fof(f3055, plain, (~ spl12_45 | ~ spl12_109 | ~ spl12_130 | spl12_168), inference(avatar_contradiction_clause, [], [f3054])).
fof(f3054, plain, ($false | (~ spl12_45 | ~ spl12_109 | ~ spl12_130 | spl12_168)), inference(subsumption_resolution, [], [f3053, f982])).
fof(f982, plain, ((e20 = op2(e21, e20)) | ~ spl12_109), inference(avatar_component_clause, [], [f980])).
fof(f980, plain, (spl12_109 <=> (e20 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl12_109])])).
fof(f3053, plain, (~ (e20 = op2(e21, e20)) | (~ spl12_45 | ~ spl12_130 | spl12_168)), inference(forward_demodulation, [], [f3052, f1266])).
fof(f3052, plain, (~ (op2(e21, e20) = h4(e10)) | (~ spl12_45 | ~ spl12_130 | spl12_168)), inference(forward_demodulation, [], [f3051, f678])).
fof(f678, plain, ((e10 = op1(e11, e10)) | ~ spl12_45), inference(avatar_component_clause, [], [f676])).
fof(f676, plain, (spl12_45 <=> (e10 = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl12_45])])).
fof(f3051, plain, (~ (op2(e21, e20) = h4(op1(e11, e10))) | (~ spl12_130 | spl12_168)), inference(forward_demodulation, [], [f1492, f1277])).
fof(f1492, plain, (~ (h4(op1(e11, e10)) = op2(h4(e11), e20)) | spl12_168), inference(avatar_component_clause, [], [f1490])).
fof(f1490, plain, (spl12_168 <=> (h4(op1(e11, e10)) = op2(h4(e11), e20))), introduced(avatar_definition, [new_symbols(naming, [spl12_168])])).
fof(f3037, plain, (~ spl12_52 | ~ spl12_116 | spl12_169), inference(avatar_contradiction_clause, [], [f3036])).
fof(f3036, plain, ($false | (~ spl12_52 | ~ spl12_116 | spl12_169)), inference(subsumption_resolution, [], [f3035, f1011])).
fof(f1011, plain, ((e23 = op2(e20, e23)) | ~ spl12_116), inference(avatar_component_clause, [], [f1009])).
fof(f1009, plain, (spl12_116 <=> (e23 = op2(e20, e23))), introduced(avatar_definition, [new_symbols(naming, [spl12_116])])).
fof(f3035, plain, (~ (e23 = op2(e20, e23)) | (~ spl12_52 | spl12_169)), inference(forward_demodulation, [], [f3034, f412])).
fof(f412, plain, (e23 = h4(e13)), inference(cnf_transformation, [], [f17])).
fof(f3034, plain, (~ (op2(e20, e23) = h4(e13)) | (~ spl12_52 | spl12_169)), inference(forward_demodulation, [], [f1496, f707])).
fof(f707, plain, ((e13 = op1(e10, e13)) | ~ spl12_52), inference(avatar_component_clause, [], [f705])).
fof(f705, plain, (spl12_52 <=> (e13 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl12_52])])).
fof(f1496, plain, (~ (op2(e20, e23) = h4(op1(e10, e13))) | spl12_169), inference(avatar_component_clause, [], [f1494])).
fof(f1494, plain, (spl12_169 <=> (op2(e20, e23) = h4(op1(e10, e13)))), introduced(avatar_definition, [new_symbols(naming, [spl12_169])])).
fof(f3022, plain, (~ spl12_38 | ~ spl12_102 | ~ spl12_130 | spl12_167), inference(avatar_contradiction_clause, [], [f3021])).
fof(f3021, plain, ($false | (~ spl12_38 | ~ spl12_102 | ~ spl12_130 | spl12_167)), inference(subsumption_resolution, [], [f3020, f952])).
fof(f952, plain, ((e21 = op2(e21, e22)) | ~ spl12_102), inference(avatar_component_clause, [], [f950])).
fof(f950, plain, (spl12_102 <=> (e21 = op2(e21, e22))), introduced(avatar_definition, [new_symbols(naming, [spl12_102])])).
fof(f3020, plain, (~ (e21 = op2(e21, e22)) | (~ spl12_38 | ~ spl12_130 | spl12_167)), inference(forward_demodulation, [], [f3019, f1277])).
fof(f3019, plain, (~ (op2(e21, e22) = h4(e11)) | (~ spl12_38 | ~ spl12_130 | spl12_167)), inference(forward_demodulation, [], [f3018, f648])).
fof(f648, plain, ((e11 = op1(e11, e12)) | ~ spl12_38), inference(avatar_component_clause, [], [f646])).
fof(f646, plain, (spl12_38 <=> (e11 = op1(e11, e12))), introduced(avatar_definition, [new_symbols(naming, [spl12_38])])).
fof(f3018, plain, (~ (op2(e21, e22) = h4(op1(e11, e12))) | (~ spl12_130 | spl12_167)), inference(forward_demodulation, [], [f1488, f1277])).
fof(f1488, plain, (~ (h4(op1(e11, e12)) = op2(h4(e11), e22)) | spl12_167), inference(avatar_component_clause, [], [f1486])).
fof(f1486, plain, (spl12_167 <=> (h4(op1(e11, e12)) = op2(h4(e11), e22))), introduced(avatar_definition, [new_symbols(naming, [spl12_167])])).
fof(f3007, plain, (~ spl12_24 | spl12_163 | ~ spl12_191), inference(avatar_contradiction_clause, [], [f3006])).
fof(f3006, plain, ($false | (~ spl12_24 | spl12_163 | ~ spl12_191)), inference(subsumption_resolution, [], [f3005, f1641])).
fof(f1641, plain, ((e23 = h3(e11)) | ~ spl12_191), inference(avatar_component_clause, [], [f1640])).
fof(f1640, plain, (spl12_191 <=> (e23 = h3(e11))), introduced(avatar_definition, [new_symbols(naming, [spl12_191])])).
fof(f3005, plain, (~ (e23 = h3(e11)) | (~ spl12_24 | spl12_163)), inference(forward_demodulation, [], [f3004, f412])).
fof(f3004, plain, (~ (h3(e11) = h4(e13)) | (~ spl12_24 | spl12_163)), inference(forward_demodulation, [], [f1472, f588])).
fof(f588, plain, ((e13 = op1(e12, e12)) | ~ spl12_24), inference(avatar_component_clause, [], [f586])).
fof(f586, plain, (spl12_24 <=> (e13 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl12_24])])).
fof(f1472, plain, (~ (h3(e11) = h4(op1(e12, e12))) | spl12_163), inference(avatar_component_clause, [], [f1470])).
fof(f1470, plain, (spl12_163 <=> (h3(e11) = h4(op1(e12, e12)))), introduced(avatar_definition, [new_symbols(naming, [spl12_163])])).
fof(f2989, plain, (~ spl12_35 | ~ spl12_99 | ~ spl12_130 | spl12_166), inference(avatar_contradiction_clause, [], [f2988])).
fof(f2988, plain, ($false | (~ spl12_35 | ~ spl12_99 | ~ spl12_130 | spl12_166)), inference(subsumption_resolution, [], [f2987, f939])).
fof(f939, plain, ((e22 = op2(e21, e23)) | ~ spl12_99), inference(avatar_component_clause, [], [f937])).
fof(f937, plain, (spl12_99 <=> (e22 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl12_99])])).
fof(f2987, plain, (~ (e22 = op2(e21, e23)) | (~ spl12_35 | ~ spl12_130 | spl12_166)), inference(forward_demodulation, [], [f2986, f1256])).
fof(f2986, plain, (~ (op2(e21, e23) = h4(e12)) | (~ spl12_35 | ~ spl12_130 | spl12_166)), inference(forward_demodulation, [], [f2985, f635])).
fof(f635, plain, ((e12 = op1(e11, e13)) | ~ spl12_35), inference(avatar_component_clause, [], [f633])).
fof(f633, plain, (spl12_35 <=> (e12 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl12_35])])).
fof(f2985, plain, (~ (op2(e21, e23) = h4(op1(e11, e13))) | (~ spl12_130 | spl12_166)), inference(forward_demodulation, [], [f1484, f1277])).
fof(f1484, plain, (~ (h4(op1(e11, e13)) = op2(h4(e11), e23)) | spl12_166), inference(avatar_component_clause, [], [f1482])).
fof(f1482, plain, (spl12_166 <=> (h4(op1(e11, e13)) = op2(h4(e11), e23))), introduced(avatar_definition, [new_symbols(naming, [spl12_166])])).
fof(f2976, plain, (~ spl12_31 | ~ spl12_95 | spl12_165), inference(avatar_contradiction_clause, [], [f2975])).
fof(f2975, plain, ($false | (~ spl12_31 | ~ spl12_95 | spl12_165)), inference(subsumption_resolution, [], [f2974, f922])).
fof(f922, plain, ((e22 = op2(e22, e20)) | ~ spl12_95), inference(avatar_component_clause, [], [f920])).
fof(f920, plain, (spl12_95 <=> (e22 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl12_95])])).
fof(f2974, plain, (~ (e22 = op2(e22, e20)) | (~ spl12_31 | spl12_165)), inference(forward_demodulation, [], [f2973, f1256])).
fof(f2973, plain, (~ (op2(e22, e20) = h4(e12)) | (~ spl12_31 | spl12_165)), inference(forward_demodulation, [], [f1480, f618])).
fof(f618, plain, ((e12 = op1(e12, e10)) | ~ spl12_31), inference(avatar_component_clause, [], [f616])).
fof(f616, plain, (spl12_31 <=> (e12 = op1(e12, e10))), introduced(avatar_definition, [new_symbols(naming, [spl12_31])])).
fof(f1480, plain, (~ (op2(e22, e20) = h4(op1(e12, e10))) | spl12_165), inference(avatar_component_clause, [], [f1478])).
fof(f1478, plain, (spl12_165 <=> (op2(e22, e20) = h4(op1(e12, e10)))), introduced(avatar_definition, [new_symbols(naming, [spl12_165])])).
fof(f2968, plain, (~ spl12_26 | ~ spl12_90 | ~ spl12_130 | spl12_164), inference(avatar_contradiction_clause, [], [f2967])).
fof(f2967, plain, ($false | (~ spl12_26 | ~ spl12_90 | ~ spl12_130 | spl12_164)), inference(subsumption_resolution, [], [f2966, f901])).
fof(f901, plain, ((e21 = op2(e22, e21)) | ~ spl12_90), inference(avatar_component_clause, [], [f899])).
fof(f899, plain, (spl12_90 <=> (e21 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl12_90])])).
fof(f2966, plain, (~ (e21 = op2(e22, e21)) | (~ spl12_26 | ~ spl12_130 | spl12_164)), inference(forward_demodulation, [], [f2965, f1277])).
fof(f2965, plain, (~ (op2(e22, e21) = h4(e11)) | (~ spl12_26 | ~ spl12_130 | spl12_164)), inference(forward_demodulation, [], [f2964, f597])).
fof(f597, plain, ((e11 = op1(e12, e11)) | ~ spl12_26), inference(avatar_component_clause, [], [f595])).
fof(f595, plain, (spl12_26 <=> (e11 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl12_26])])).
fof(f2964, plain, (~ (op2(e22, e21) = h4(op1(e12, e11))) | (~ spl12_130 | spl12_164)), inference(forward_demodulation, [], [f1476, f1277])).
fof(f1476, plain, (~ (h4(op1(e12, e11)) = op2(e22, h4(e11))) | spl12_164), inference(avatar_component_clause, [], [f1474])).
fof(f1474, plain, (spl12_164 <=> (h4(op1(e12, e11)) = op2(e22, h4(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl12_164])])).
fof(f2944, plain, (~ spl12_9 | spl12_160), inference(avatar_contradiction_clause, [], [f2943])).
fof(f2943, plain, ($false | (~ spl12_9 | spl12_160)), inference(subsumption_resolution, [], [f2942, f1266])).
fof(f2942, plain, (~ (e20 = h4(e10)) | (~ spl12_9 | spl12_160)), inference(forward_demodulation, [], [f1460, f525])).
fof(f525, plain, ((e10 = op1(e13, e11)) | ~ spl12_9), inference(avatar_component_clause, [], [f523])).
fof(f523, plain, (spl12_9 <=> (e10 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl12_9])])).
fof(f1460, plain, (~ (e20 = h4(op1(e13, e11))) | spl12_160), inference(avatar_component_clause, [], [f1458])).
fof(f1458, plain, (spl12_160 <=> (e20 = h4(op1(e13, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl12_160])])).
fof(f2936, plain, (~ spl12_29 | ~ spl12_31), inference(avatar_contradiction_clause, [], [f2935])).
fof(f2935, plain, ($false | (~ spl12_29 | ~ spl12_31)), inference(subsumption_resolution, [], [f2934, f239])).
fof(f239, plain, ~ (e10 = e12), inference(cnf_transformation, [], [f7])).
fof(f7, plain, (~ (e12 = e13) & ~ (e11 = e13) & ~ (e11 = e12) & ~ (e10 = e13) & ~ (e10 = e12) & ~ (e10 = e11)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG111+1.p', ax7)).
fof(f2934, plain, ((e10 = e12) | (~ spl12_29 | ~ spl12_31)), inference(backward_demodulation, [], [f618, f610])).
fof(f610, plain, ((e10 = op1(e12, e10)) | ~ spl12_29), inference(avatar_component_clause, [], [f608])).
fof(f608, plain, (spl12_29 <=> (e10 = op1(e12, e10))), introduced(avatar_definition, [new_symbols(naming, [spl12_29])])).
fof(f2893, plain, (~ spl12_17 | ~ spl12_81 | spl12_162), inference(avatar_contradiction_clause, [], [f2892])).
fof(f2892, plain, ($false | (~ spl12_17 | ~ spl12_81 | spl12_162)), inference(subsumption_resolution, [], [f2891, f863])).
fof(f863, plain, ((e20 = op2(e22, e23)) | ~ spl12_81), inference(avatar_component_clause, [], [f861])).
fof(f861, plain, (spl12_81 <=> (e20 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl12_81])])).
fof(f2891, plain, (~ (e20 = op2(e22, e23)) | (~ spl12_17 | spl12_162)), inference(forward_demodulation, [], [f2890, f1266])).
fof(f2890, plain, (~ (op2(e22, e23) = h4(e10)) | (~ spl12_17 | spl12_162)), inference(forward_demodulation, [], [f1468, f559])).
fof(f559, plain, ((e10 = op1(e12, e13)) | ~ spl12_17), inference(avatar_component_clause, [], [f557])).
fof(f557, plain, (spl12_17 <=> (e10 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl12_17])])).
fof(f1468, plain, (~ (op2(e22, e23) = h4(op1(e12, e13))) | spl12_162), inference(avatar_component_clause, [], [f1466])).
fof(f1466, plain, (spl12_162 <=> (op2(e22, e23) = h4(op1(e12, e13)))), introduced(avatar_definition, [new_symbols(naming, [spl12_162])])).
fof(f2878, plain, (~ spl12_16 | ~ spl12_80 | spl12_161), inference(avatar_contradiction_clause, [], [f2877])).
fof(f2877, plain, ($false | (~ spl12_16 | ~ spl12_80 | spl12_161)), inference(subsumption_resolution, [], [f2876, f858])).
fof(f858, plain, ((e23 = op2(e23, e20)) | ~ spl12_80), inference(avatar_component_clause, [], [f856])).
fof(f856, plain, (spl12_80 <=> (e23 = op2(e23, e20))), introduced(avatar_definition, [new_symbols(naming, [spl12_80])])).
fof(f2876, plain, (~ (e23 = op2(e23, e20)) | (~ spl12_16 | spl12_161)), inference(forward_demodulation, [], [f2875, f412])).
fof(f2875, plain, (~ (op2(e23, e20) = h4(e13)) | (~ spl12_16 | spl12_161)), inference(forward_demodulation, [], [f1464, f554])).
fof(f554, plain, ((e13 = op1(e13, e10)) | ~ spl12_16), inference(avatar_component_clause, [], [f552])).
fof(f552, plain, (spl12_16 <=> (e13 = op1(e13, e10))), introduced(avatar_definition, [new_symbols(naming, [spl12_16])])).
fof(f1464, plain, (~ (op2(e23, e20) = h4(op1(e13, e10))) | spl12_161), inference(avatar_component_clause, [], [f1462])).
fof(f1462, plain, (spl12_161 <=> (op2(e23, e20) = h4(op1(e13, e10)))), introduced(avatar_definition, [new_symbols(naming, [spl12_161])])).
fof(f2859, plain, (~ spl12_7 | ~ spl12_71 | spl12_159), inference(avatar_contradiction_clause, [], [f2858])).
fof(f2858, plain, ($false | (~ spl12_7 | ~ spl12_71 | spl12_159)), inference(subsumption_resolution, [], [f2857, f820])).
fof(f820, plain, ((e22 = op2(e23, e22)) | ~ spl12_71), inference(avatar_component_clause, [], [f818])).
fof(f818, plain, (spl12_71 <=> (e22 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl12_71])])).
fof(f2857, plain, (~ (e22 = op2(e23, e22)) | (~ spl12_7 | spl12_159)), inference(forward_demodulation, [], [f2856, f1256])).
fof(f2856, plain, (~ (op2(e23, e22) = h4(e12)) | (~ spl12_7 | spl12_159)), inference(forward_demodulation, [], [f1456, f516])).
fof(f516, plain, ((e12 = op1(e13, e12)) | ~ spl12_7), inference(avatar_component_clause, [], [f514])).
fof(f514, plain, (spl12_7 <=> (e12 = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl12_7])])).
fof(f1456, plain, (~ (op2(e23, e22) = h4(op1(e13, e12))) | spl12_159), inference(avatar_component_clause, [], [f1454])).
fof(f1454, plain, (spl12_159 <=> (op2(e23, e22) = h4(op1(e13, e12)))), introduced(avatar_definition, [new_symbols(naming, [spl12_159])])).
fof(f2837, plain, (~ spl12_2 | spl12_158), inference(avatar_contradiction_clause, [], [f2836])).
fof(f2836, plain, ($false | (~ spl12_2 | spl12_158)), inference(trivial_inequality_removal, [], [f2835])).
fof(f2835, plain, (~ (h4(e11) = h4(e11)) | (~ spl12_2 | spl12_158)), inference(forward_demodulation, [], [f1452, f495])).
fof(f495, plain, ((e11 = op1(e13, e13)) | ~ spl12_2), inference(avatar_component_clause, [], [f493])).
fof(f493, plain, (spl12_2 <=> (e11 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl12_2])])).
fof(f1452, plain, (~ (h4(e11) = h4(op1(e13, e13))) | spl12_158), inference(avatar_component_clause, [], [f1450])).
fof(f1450, plain, (spl12_158 <=> (h4(e11) = h4(op1(e13, e13)))), introduced(avatar_definition, [new_symbols(naming, [spl12_158])])).
fof(f2815, plain, (~ spl12_95 | ~ spl12_111), inference(avatar_contradiction_clause, [], [f2814])).
fof(f2814, plain, ($false | (~ spl12_95 | ~ spl12_111)), inference(subsumption_resolution, [], [f2813, f990])).
fof(f990, plain, ((e22 = op2(e21, e20)) | ~ spl12_111), inference(avatar_component_clause, [], [f988])).
fof(f988, plain, (spl12_111 <=> (e22 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl12_111])])).
fof(f2813, plain, (~ (e22 = op2(e21, e20)) | ~ spl12_95), inference(forward_demodulation, [], [f193, f922])).
fof(f193, plain, ~ (op2(e21, e20) = op2(e22, e20)), inference(cnf_transformation, [], [f6])).
fof(f6, plain, (~ (op2(e23, e22) = op2(e23, e23)) & ~ (op2(e23, e21) = op2(e23, e23)) & ~ (op2(e23, e21) = op2(e23, e22)) & ~ (op2(e23, e20) = op2(e23, e23)) & ~ (op2(e23, e20) = op2(e23, e22)) & ~ (op2(e23, e20) = op2(e23, e21)) & ~ (op2(e22, e22) = op2(e22, e23)) & ~ (op2(e22, e21) = op2(e22, e23)) & ~ (op2(e22, e21) = op2(e22, e22)) & ~ (op2(e22, e20) = op2(e22, e23)) & ~ (op2(e22, e20) = op2(e22, e22)) & ~ (op2(e22, e20) = op2(e22, e21)) & ~ (op2(e21, e22) = op2(e21, e23)) & ~ (op2(e21, e21) = op2(e21, e23)) & ~ (op2(e21, e21) = op2(e21, e22)) & ~ (op2(e21, e20) = op2(e21, e23)) & ~ (op2(e21, e20) = op2(e21, e22)) & ~ (op2(e21, e20) = op2(e21, e21)) & ~ (op2(e20, e22) = op2(e20, e23)) & ~ (op2(e20, e21) = op2(e20, e23)) & ~ (op2(e20, e21) = op2(e20, e22)) & ~ (op2(e20, e20) = op2(e20, e23)) & ~ (op2(e20, e20) = op2(e20, e22)) & ~ (op2(e20, e20) = op2(e20, e21)) & ~ (op2(e22, e23) = op2(e23, e23)) & ~ (op2(e21, e23) = op2(e23, e23)) & ~ (op2(e21, e23) = op2(e22, e23)) & ~ (op2(e20, e23) = op2(e23, e23)) & ~ (op2(e20, e23) = op2(e22, e23)) & ~ (op2(e20, e23) = op2(e21, e23)) & ~ (op2(e22, e22) = op2(e23, e22)) & ~ (op2(e21, e22) = op2(e23, e22)) & ~ (op2(e21, e22) = op2(e22, e22)) & ~ (op2(e20, e22) = op2(e23, e22)) & ~ (op2(e20, e22) = op2(e22, e22)) & ~ (op2(e20, e22) = op2(e21, e22)) & ~ (op2(e22, e21) = op2(e23, e21)) & ~ (op2(e21, e21) = op2(e23, e21)) & ~ (op2(e21, e21) = op2(e22, e21)) & ~ (op2(e20, e21) = op2(e23, e21)) & ~ (op2(e20, e21) = op2(e22, e21)) & ~ (op2(e20, e21) = op2(e21, e21)) & ~ (op2(e22, e20) = op2(e23, e20)) & ~ (op2(e21, e20) = op2(e23, e20)) & ~ (op2(e21, e20) = op2(e22, e20)) & ~ (op2(e20, e20) = op2(e23, e20)) & ~ (op2(e20, e20) = op2(e22, e20)) & ~ (op2(e20, e20) = op2(e21, e20))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG111+1.p', ax6)).
fof(f2802, plain, (~ spl12_93 | ~ spl12_95), inference(avatar_contradiction_clause, [], [f2801])).
fof(f2801, plain, ($false | (~ spl12_93 | ~ spl12_95)), inference(subsumption_resolution, [], [f2800, f245])).
fof(f245, plain, ~ (e20 = e22), inference(cnf_transformation, [], [f8])).
fof(f8, plain, (~ (e22 = e23) & ~ (e21 = e23) & ~ (e21 = e22) & ~ (e20 = e23) & ~ (e20 = e22) & ~ (e20 = e21)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG111+1.p', ax8)).
fof(f2800, plain, ((e20 = e22) | (~ spl12_93 | ~ spl12_95)), inference(forward_demodulation, [], [f922, f914])).
fof(f914, plain, ((e20 = op2(e22, e20)) | ~ spl12_93), inference(avatar_component_clause, [], [f912])).
fof(f912, plain, (spl12_93 <=> (e20 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl12_93])])).
fof(f2767, plain, (~ spl12_210 | ~ spl12_44 | ~ spl12_130 | spl12_157), inference(avatar_split_clause, [], [f2436, f1446, f1276, f671, f1754])).
fof(f1754, plain, (spl12_210 <=> (e23 = h2(e11))), introduced(avatar_definition, [new_symbols(naming, [spl12_210])])).
fof(f671, plain, (spl12_44 <=> (e13 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl12_44])])).
fof(f1446, plain, (spl12_157 <=> (h4(op1(e11, e11)) = op2(h4(e11), h4(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl12_157])])).
fof(f2436, plain, (~ (e23 = h2(e11)) | (~ spl12_44 | ~ spl12_130 | spl12_157)), inference(forward_demodulation, [], [f2434, f412])).
fof(f2434, plain, (~ (h2(e11) = h4(e13)) | (~ spl12_44 | ~ spl12_130 | spl12_157)), inference(backward_demodulation, [], [f2239, f673])).
fof(f673, plain, ((e13 = op1(e11, e11)) | ~ spl12_44), inference(avatar_component_clause, [], [f671])).
fof(f2239, plain, (~ (h2(e11) = h4(op1(e11, e11))) | (~ spl12_130 | spl12_157)), inference(forward_demodulation, [], [f2231, f406])).
fof(f406, plain, (op2(e21, e21) = h2(e11)), inference(cnf_transformation, [], [f15])).
fof(f15, plain, ((h2(e12) = op2(op2(e21, op2(e21, e21)), op2(e21, e21))) & (op2(e21, e21) = h2(e11)) & (h2(e10) = op2(e21, op2(e21, e21))) & (e21 = h2(e13))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG111+1.p', ax15)).
fof(f2231, plain, (~ (op2(e21, e21) = h4(op1(e11, e11))) | (~ spl12_130 | spl12_157)), inference(backward_demodulation, [], [f1448, f1277])).
fof(f1448, plain, (~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | spl12_157), inference(avatar_component_clause, [], [f1446])).
fof(f2711, plain, (~ spl12_80 | ~ spl12_96), inference(avatar_split_clause, [], [f2709, f924, f856])).
fof(f924, plain, (spl12_96 <=> (e23 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl12_96])])).
fof(f2709, plain, (~ (e23 = op2(e23, e20)) | ~ spl12_96), inference(backward_demodulation, [], [f195, f926])).
fof(f926, plain, ((e23 = op2(e22, e20)) | ~ spl12_96), inference(avatar_component_clause, [], [f924])).
fof(f195, plain, ~ (op2(e22, e20) = op2(e23, e20)), inference(cnf_transformation, [], [f6])).
fof(f2701, plain, (~ spl12_102 | ~ spl12_103), inference(avatar_contradiction_clause, [], [f2700])).
fof(f2700, plain, ($false | (~ spl12_102 | ~ spl12_103)), inference(subsumption_resolution, [], [f2699, f247])).
fof(f247, plain, ~ (e21 = e22), inference(cnf_transformation, [], [f8])).
fof(f2699, plain, ((e21 = e22) | (~ spl12_102 | ~ spl12_103)), inference(backward_demodulation, [], [f956, f952])).
fof(f956, plain, ((e22 = op2(e21, e22)) | ~ spl12_103), inference(avatar_component_clause, [], [f954])).
fof(f954, plain, (spl12_103 <=> (e22 = op2(e21, e22))), introduced(avatar_definition, [new_symbols(naming, [spl12_103])])).
fof(f2680, plain, (~ spl12_123 | ~ spl12_124), inference(avatar_contradiction_clause, [], [f2679])).
fof(f2679, plain, ($false | (~ spl12_123 | ~ spl12_124)), inference(subsumption_resolution, [], [f2678, f249])).
fof(f249, plain, ~ (e22 = e23), inference(cnf_transformation, [], [f8])).
fof(f2678, plain, ((e22 = e23) | (~ spl12_123 | ~ spl12_124)), inference(forward_demodulation, [], [f1045, f1041])).
fof(f1041, plain, ((e22 = op2(e20, e21)) | ~ spl12_123), inference(avatar_component_clause, [], [f1039])).
fof(f1039, plain, (spl12_123 <=> (e22 = op2(e20, e21))), introduced(avatar_definition, [new_symbols(naming, [spl12_123])])).
fof(f1045, plain, ((e23 = op2(e20, e21)) | ~ spl12_124), inference(avatar_component_clause, [], [f1043])).
fof(f1043, plain, (spl12_124 <=> (e23 = op2(e20, e21))), introduced(avatar_definition, [new_symbols(naming, [spl12_124])])).
fof(f2658, plain, (~ spl12_110 | ~ spl12_155), inference(avatar_split_clause, [], [f2650, f1407, f984])).
fof(f984, plain, (spl12_110 <=> (e21 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl12_110])])).
fof(f2650, plain, (~ (e21 = op2(e21, e20)) | ~ spl12_155), inference(backward_demodulation, [], [f1228, f1408])).
fof(f1228, plain, ~ (op2(e21, e20) = h1(e11)), inference(backward_demodulation, [], [f190, f402])).
fof(f402, plain, (op2(e20, e20) = h1(e11)), inference(cnf_transformation, [], [f14])).
fof(f14, plain, ((h1(e12) = op2(op2(e20, op2(e20, e20)), op2(e20, e20))) & (op2(e20, e20) = h1(e11)) & (h1(e10) = op2(e20, op2(e20, e20))) & (e20 = h1(e13))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG111+1.p', ax14)).
fof(f190, plain, ~ (op2(e20, e20) = op2(e21, e20)), inference(cnf_transformation, [], [f6])).
fof(f2633, plain, (~ spl12_114 | ~ spl12_130), inference(avatar_split_clause, [], [f2503, f1276, f1001])).
fof(f1001, plain, (spl12_114 <=> (e21 = op2(e20, e23))), introduced(avatar_definition, [new_symbols(naming, [spl12_114])])).
fof(f2503, plain, (~ (e21 = op2(e20, e23)) | ~ spl12_130), inference(forward_demodulation, [], [f1257, f1277])).
fof(f1257, plain, ~ (op2(e20, e23) = h4(e11)), inference(backward_demodulation, [], [f210, f414])).
fof(f210, plain, ~ (op2(e20, e23) = op2(e23, e23)), inference(cnf_transformation, [], [f6])).
fof(f2621, plain, (~ spl12_89 | ~ spl12_73), inference(avatar_split_clause, [], [f2620, f827, f895])).
fof(f895, plain, (spl12_89 <=> (e20 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl12_89])])).
fof(f827, plain, (spl12_73 <=> (e20 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl12_73])])).
fof(f2620, plain, (~ (e20 = op2(e22, e21)) | ~ spl12_73), inference(forward_demodulation, [], [f201, f829])).
fof(f829, plain, ((e20 = op2(e23, e21)) | ~ spl12_73), inference(avatar_component_clause, [], [f827])).
fof(f201, plain, ~ (op2(e22, e21) = op2(e23, e21)), inference(cnf_transformation, [], [f6])).
fof(f2584, plain, (~ spl12_45 | ~ spl12_46), inference(avatar_contradiction_clause, [], [f2583])).
fof(f2583, plain, ($false | (~ spl12_45 | ~ spl12_46)), inference(subsumption_resolution, [], [f2582, f238])).
fof(f238, plain, ~ (e10 = e11), inference(cnf_transformation, [], [f7])).
fof(f2582, plain, ((e10 = e11) | (~ spl12_45 | ~ spl12_46)), inference(forward_demodulation, [], [f682, f678])).
fof(f682, plain, ((e11 = op1(e11, e10)) | ~ spl12_46), inference(avatar_component_clause, [], [f680])).
fof(f680, plain, (spl12_46 <=> (e11 = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl12_46])])).
fof(f2564, plain, (~ spl12_90 | ~ spl12_92), inference(avatar_contradiction_clause, [], [f2563])).
fof(f2563, plain, ($false | (~ spl12_90 | ~ spl12_92)), inference(subsumption_resolution, [], [f2562, f248])).
fof(f248, plain, ~ (e21 = e23), inference(cnf_transformation, [], [f8])).
fof(f2562, plain, ((e21 = e23) | (~ spl12_90 | ~ spl12_92)), inference(backward_demodulation, [], [f909, f901])).
fof(f909, plain, ((e23 = op2(e22, e21)) | ~ spl12_92), inference(avatar_component_clause, [], [f907])).
fof(f907, plain, (spl12_92 <=> (e23 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl12_92])])).
fof(f2513, plain, (~ spl12_143 | ~ spl12_123), inference(avatar_split_clause, [], [f2392, f1039, f1346])).
fof(f1346, plain, (spl12_143 <=> (e22 = h2(e11))), introduced(avatar_definition, [new_symbols(naming, [spl12_143])])).
fof(f2392, plain, (~ (e22 = h2(e11)) | ~ spl12_123), inference(backward_demodulation, [], [f1237, f1041])).
fof(f1237, plain, ~ (op2(e20, e21) = h2(e11)), inference(backward_demodulation, [], [f196, f406])).
fof(f196, plain, ~ (op2(e20, e21) = op2(e21, e21)), inference(cnf_transformation, [], [f6])).
fof(f2500, plain, (~ spl12_115 | ~ spl12_123), inference(avatar_split_clause, [], [f2499, f1039, f1005])).
fof(f1005, plain, (spl12_115 <=> (e22 = op2(e20, e23))), introduced(avatar_definition, [new_symbols(naming, [spl12_115])])).
fof(f2499, plain, (~ (e22 = op2(e20, e23)) | ~ spl12_123), inference(forward_demodulation, [], [f218, f1041])).
fof(f218, plain, ~ (op2(e20, e21) = op2(e20, e23)), inference(cnf_transformation, [], [f6])).
fof(f2482, plain, (~ spl12_77 | ~ spl12_73), inference(avatar_split_clause, [], [f2363, f827, f844])).
fof(f844, plain, (spl12_77 <=> (e20 = op2(e23, e20))), introduced(avatar_definition, [new_symbols(naming, [spl12_77])])).
fof(f2363, plain, (~ (e20 = op2(e23, e20)) | ~ spl12_73), inference(forward_demodulation, [], [f232, f829])).
fof(f232, plain, ~ (op2(e23, e20) = op2(e23, e21)), inference(cnf_transformation, [], [f6])).
fof(f2479, plain, (~ spl12_64 | ~ spl12_16), inference(avatar_split_clause, [], [f2478, f552, f756])).
fof(f756, plain, (spl12_64 <=> (op1(e10, e10) = e13)), introduced(avatar_definition, [new_symbols(naming, [spl12_64])])).
fof(f2478, plain, (~ (op1(e10, e10) = e13) | ~ spl12_16), inference(forward_demodulation, [], [f144, f554])).
fof(f144, plain, ~ (op1(e10, e10) = op1(e13, e10)), inference(cnf_transformation, [], [f5])).
fof(f5, plain, (~ (op1(e13, e12) = op1(e13, e13)) & ~ (op1(e13, e11) = op1(e13, e13)) & ~ (op1(e13, e11) = op1(e13, e12)) & ~ (op1(e13, e10) = op1(e13, e13)) & ~ (op1(e13, e10) = op1(e13, e12)) & ~ (op1(e13, e10) = op1(e13, e11)) & ~ (op1(e12, e12) = op1(e12, e13)) & ~ (op1(e12, e11) = op1(e12, e13)) & ~ (op1(e12, e11) = op1(e12, e12)) & ~ (op1(e12, e10) = op1(e12, e13)) & ~ (op1(e12, e10) = op1(e12, e12)) & ~ (op1(e12, e10) = op1(e12, e11)) & ~ (op1(e11, e12) = op1(e11, e13)) & ~ (op1(e11, e11) = op1(e11, e13)) & ~ (op1(e11, e11) = op1(e11, e12)) & ~ (op1(e11, e10) = op1(e11, e13)) & ~ (op1(e11, e10) = op1(e11, e12)) & ~ (op1(e11, e10) = op1(e11, e11)) & ~ (op1(e10, e12) = op1(e10, e13)) & ~ (op1(e10, e11) = op1(e10, e13)) & ~ (op1(e10, e11) = op1(e10, e12)) & ~ (op1(e10, e10) = op1(e10, e13)) & ~ (op1(e10, e10) = op1(e10, e12)) & ~ (op1(e10, e10) = op1(e10, e11)) & ~ (op1(e12, e13) = op1(e13, e13)) & ~ (op1(e11, e13) = op1(e13, e13)) & ~ (op1(e11, e13) = op1(e12, e13)) & ~ (op1(e10, e13) = op1(e13, e13)) & ~ (op1(e10, e13) = op1(e12, e13)) & ~ (op1(e10, e13) = op1(e11, e13)) & ~ (op1(e12, e12) = op1(e13, e12)) & ~ (op1(e11, e12) = op1(e13, e12)) & ~ (op1(e11, e12) = op1(e12, e12)) & ~ (op1(e10, e12) = op1(e13, e12)) & ~ (op1(e10, e12) = op1(e12, e12)) & ~ (op1(e10, e12) = op1(e11, e12)) & ~ (op1(e12, e11) = op1(e13, e11)) & ~ (op1(e11, e11) = op1(e13, e11)) & ~ (op1(e11, e11) = op1(e12, e11)) & ~ (op1(e10, e11) = op1(e13, e11)) & ~ (op1(e10, e11) = op1(e12, e11)) & ~ (op1(e10, e11) = op1(e11, e11)) & ~ (op1(e12, e10) = op1(e13, e10)) & ~ (op1(e11, e10) = op1(e13, e10)) & ~ (op1(e11, e10) = op1(e12, e10)) & ~ (op1(e10, e10) = op1(e13, e10)) & ~ (op1(e10, e10) = op1(e12, e10)) & ~ (op1(e10, e10) = op1(e11, e10))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG111+1.p', ax5)).
fof(f2472, plain, (~ spl12_50 | ~ spl12_2), inference(avatar_split_clause, [], [f2351, f493, f697])).
fof(f697, plain, (spl12_50 <=> (e11 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl12_50])])).
fof(f2351, plain, (~ (e11 = op1(e10, e13)) | ~ spl12_2), inference(forward_demodulation, [], [f162, f495])).
fof(f162, plain, ~ (op1(e10, e13) = op1(e13, e13)), inference(cnf_transformation, [], [f5])).
fof(f2461, plain, (~ spl12_25 | ~ spl12_9), inference(avatar_split_clause, [], [f2460, f523, f591])).
fof(f591, plain, (spl12_25 <=> (e10 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl12_25])])).
fof(f2460, plain, (~ (e10 = op1(e12, e11)) | ~ spl12_9), inference(forward_demodulation, [], [f153, f525])).
fof(f153, plain, ~ (op1(e12, e11) = op1(e13, e11)), inference(cnf_transformation, [], [f5])).
fof(f2430, plain, (~ spl12_53 | ~ spl12_56), inference(avatar_contradiction_clause, [], [f2429])).
fof(f2429, plain, ($false | (~ spl12_53 | ~ spl12_56)), inference(subsumption_resolution, [], [f2427, f240])).
fof(f240, plain, ~ (e10 = e13), inference(cnf_transformation, [], [f7])).
fof(f2427, plain, ((e10 = e13) | (~ spl12_53 | ~ spl12_56)), inference(backward_demodulation, [], [f724, f712])).
fof(f724, plain, ((e13 = op1(e10, e12)) | ~ spl12_56), inference(avatar_component_clause, [], [f722])).
fof(f722, plain, (spl12_56 <=> (e13 = op1(e10, e12))), introduced(avatar_definition, [new_symbols(naming, [spl12_56])])).
fof(f2419, plain, (~ spl12_73 | ~ spl12_74), inference(avatar_contradiction_clause, [], [f2418])).
fof(f2418, plain, ($false | (~ spl12_73 | ~ spl12_74)), inference(subsumption_resolution, [], [f2417, f244])).
fof(f244, plain, ~ (e20 = e21), inference(cnf_transformation, [], [f8])).
fof(f2417, plain, ((e20 = e21) | (~ spl12_73 | ~ spl12_74)), inference(forward_demodulation, [], [f833, f829])).
fof(f833, plain, ((e21 = op2(e23, e21)) | ~ spl12_74), inference(avatar_component_clause, [], [f831])).
fof(f831, plain, (spl12_74 <=> (e21 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl12_74])])).
fof(f2414, plain, (~ spl12_88 | spl12_191), inference(avatar_contradiction_clause, [], [f2413])).
fof(f2413, plain, ($false | (~ spl12_88 | spl12_191)), inference(subsumption_resolution, [], [f2412, f1642])).
fof(f1642, plain, (~ (e23 = h3(e11)) | spl12_191), inference(avatar_component_clause, [], [f1640])).
fof(f2412, plain, ((e23 = h3(e11)) | ~ spl12_88), inference(backward_demodulation, [], [f410, f892])).
fof(f892, plain, ((e23 = op2(e22, e22)) | ~ spl12_88), inference(avatar_component_clause, [], [f890])).
fof(f890, plain, (spl12_88 <=> (e23 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl12_88])])).
fof(f410, plain, (op2(e22, e22) = h3(e11)), inference(cnf_transformation, [], [f16])).
fof(f16, plain, ((h3(e12) = op2(op2(e22, op2(e22, e22)), op2(e22, e22))) & (op2(e22, e22) = h3(e11)) & (h3(e10) = op2(e22, op2(e22, e22))) & (e22 = h3(e13))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG111+1.p', ax16)).
fof(f2406, plain, (~ spl12_108 | spl12_210), inference(avatar_contradiction_clause, [], [f2405])).
fof(f2405, plain, ($false | (~ spl12_108 | spl12_210)), inference(subsumption_resolution, [], [f2404, f1756])).
fof(f1756, plain, (~ (e23 = h2(e11)) | spl12_210), inference(avatar_component_clause, [], [f1754])).
fof(f2404, plain, ((e23 = h2(e11)) | ~ spl12_108), inference(backward_demodulation, [], [f406, f977])).
fof(f977, plain, ((e23 = op2(e21, e21)) | ~ spl12_108), inference(avatar_component_clause, [], [f975])).
fof(f975, plain, (spl12_108 <=> (e23 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl12_108])])).
fof(f2400, plain, (~ spl12_117 | ~ spl12_118), inference(avatar_contradiction_clause, [], [f2399])).
fof(f2399, plain, ($false | (~ spl12_117 | ~ spl12_118)), inference(subsumption_resolution, [], [f2398, f244])).
fof(f2398, plain, ((e20 = e21) | (~ spl12_117 | ~ spl12_118)), inference(backward_demodulation, [], [f1020, f1016])).
fof(f1020, plain, ((e21 = op2(e20, e22)) | ~ spl12_118), inference(avatar_component_clause, [], [f1018])).
fof(f1018, plain, (spl12_118 <=> (e21 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl12_118])])).
fof(f2384, plain, (~ spl12_123 | ~ spl12_119), inference(avatar_split_clause, [], [f2383, f1022, f1039])).
fof(f1022, plain, (spl12_119 <=> (e22 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl12_119])])).
fof(f2383, plain, (~ (e22 = op2(e20, e21)) | ~ spl12_119), inference(forward_demodulation, [], [f217, f1024])).
fof(f1024, plain, ((e22 = op2(e20, e22)) | ~ spl12_119), inference(avatar_component_clause, [], [f1022])).
fof(f217, plain, ~ (op2(e20, e21) = op2(e20, e22)), inference(cnf_transformation, [], [f6])).
fof(f2380, plain, (spl12_123 | ~ spl12_130), inference(avatar_split_clause, [], [f2230, f1276, f1039])).
fof(f2230, plain, ((e22 = op2(e20, e21)) | ~ spl12_130), inference(backward_demodulation, [], [f1264, f1277])).
fof(f1264, plain, (e22 = op2(e20, h4(e11))), inference(backward_demodulation, [], [f1227, f414])).
fof(f2379, plain, (~ spl12_98 | ~ spl12_130), inference(avatar_split_clause, [], [f2378, f1276, f933])).
fof(f933, plain, (spl12_98 <=> (e21 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl12_98])])).
fof(f2378, plain, (~ (e21 = op2(e21, e23)) | ~ spl12_130), inference(forward_demodulation, [], [f1258, f1277])).
fof(f1258, plain, ~ (op2(e21, e23) = h4(e11)), inference(backward_demodulation, [], [f212, f414])).
fof(f212, plain, ~ (op2(e21, e23) = op2(e23, e23)), inference(cnf_transformation, [], [f6])).
fof(f2350, plain, (~ spl12_51 | ~ spl12_59), inference(avatar_split_clause, [], [f2349, f735, f701])).
fof(f701, plain, (spl12_51 <=> (e12 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl12_51])])).
fof(f2349, plain, (~ (e12 = op1(e10, e13)) | ~ spl12_59), inference(forward_demodulation, [], [f170, f737])).
fof(f170, plain, ~ (op1(e10, e11) = op1(e10, e13)), inference(cnf_transformation, [], [f5])).
fof(f2347, plain, (~ spl12_44 | ~ spl12_40), inference(avatar_split_clause, [], [f2346, f654, f671])).
fof(f654, plain, (spl12_40 <=> (e13 = op1(e11, e12))), introduced(avatar_definition, [new_symbols(naming, [spl12_40])])).
fof(f2346, plain, (~ (e13 = op1(e11, e11)) | ~ spl12_40), inference(forward_demodulation, [], [f175, f656])).
fof(f656, plain, ((e13 = op1(e11, e12)) | ~ spl12_40), inference(avatar_component_clause, [], [f654])).
fof(f175, plain, ~ (op1(e11, e11) = op1(e11, e12)), inference(cnf_transformation, [], [f5])).
fof(f2345, plain, (~ spl12_34 | ~ spl12_2), inference(avatar_split_clause, [], [f2180, f493, f629])).
fof(f629, plain, (spl12_34 <=> (e11 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl12_34])])).
fof(f2180, plain, (~ (e11 = op1(e11, e13)) | ~ spl12_2), inference(forward_demodulation, [], [f164, f495])).
fof(f164, plain, ~ (op1(e11, e13) = op1(e13, e13)), inference(cnf_transformation, [], [f5])).
fof(f2326, plain, (~ spl12_9 | ~ spl12_10), inference(avatar_contradiction_clause, [], [f2325])).
fof(f2325, plain, ($false | (~ spl12_9 | ~ spl12_10)), inference(subsumption_resolution, [], [f2324, f238])).
fof(f2324, plain, ((e10 = e11) | (~ spl12_9 | ~ spl12_10)), inference(backward_demodulation, [], [f529, f525])).
fof(f529, plain, ((e11 = op1(e13, e11)) | ~ spl12_10), inference(avatar_component_clause, [], [f527])).
fof(f527, plain, (spl12_10 <=> (e11 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl12_10])])).
fof(f2316, plain, (~ spl12_15 | ~ spl12_31), inference(avatar_split_clause, [], [f2314, f616, f548])).
fof(f548, plain, (spl12_15 <=> (e12 = op1(e13, e10))), introduced(avatar_definition, [new_symbols(naming, [spl12_15])])).
fof(f2314, plain, (~ (e12 = op1(e13, e10)) | ~ spl12_31), inference(backward_demodulation, [], [f147, f618])).
fof(f147, plain, ~ (op1(e12, e10) = op1(e13, e10)), inference(cnf_transformation, [], [f5])).
fof(f2295, plain, (~ spl12_43 | ~ spl12_59), inference(avatar_split_clause, [], [f2287, f735, f667])).
fof(f667, plain, (spl12_43 <=> (e12 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl12_43])])).
fof(f2287, plain, (~ (e12 = op1(e11, e11)) | ~ spl12_59), inference(backward_demodulation, [], [f148, f737])).
fof(f148, plain, ~ (op1(e10, e11) = op1(e11, e11)), inference(cnf_transformation, [], [f5])).
fof(f2256, plain, (~ spl12_107 | spl12_143), inference(avatar_contradiction_clause, [], [f2255])).
fof(f2255, plain, ($false | (~ spl12_107 | spl12_143)), inference(subsumption_resolution, [], [f2254, f1348])).
fof(f1348, plain, (~ (e22 = h2(e11)) | spl12_143), inference(avatar_component_clause, [], [f1346])).
fof(f2254, plain, ((e22 = h2(e11)) | ~ spl12_107), inference(backward_demodulation, [], [f406, f973])).
fof(f973, plain, ((e22 = op2(e21, e21)) | ~ spl12_107), inference(avatar_component_clause, [], [f971])).
fof(f971, plain, (spl12_107 <=> (e22 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl12_107])])).
fof(f2238, plain, (~ spl12_122 | ~ spl12_130), inference(avatar_contradiction_clause, [], [f2237])).
fof(f2237, plain, ($false | (~ spl12_122 | ~ spl12_130)), inference(subsumption_resolution, [], [f2236, f247])).
fof(f2236, plain, ((e21 = e22) | (~ spl12_122 | ~ spl12_130)), inference(forward_demodulation, [], [f2230, f1037])).
fof(f1037, plain, ((e21 = op2(e20, e21)) | ~ spl12_122), inference(avatar_component_clause, [], [f1035])).
fof(f1035, plain, (spl12_122 <=> (e21 = op2(e20, e21))), introduced(avatar_definition, [new_symbols(naming, [spl12_122])])).
fof(f2235, plain, (spl12_73 | ~ spl12_130), inference(avatar_split_clause, [], [f2229, f1276, f827])).
fof(f2229, plain, ((e20 = op2(e23, e21)) | ~ spl12_130), inference(backward_demodulation, [], [f1263, f1277])).
fof(f2201, plain, (~ spl12_59 | ~ spl12_11), inference(avatar_split_clause, [], [f2200, f531, f735])).
fof(f531, plain, (spl12_11 <=> (e12 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl12_11])])).
fof(f2200, plain, (~ (e12 = op1(e10, e11)) | ~ spl12_11), inference(forward_demodulation, [], [f150, f533])).
fof(f533, plain, ((e12 = op1(e13, e11)) | ~ spl12_11), inference(avatar_component_clause, [], [f531])).
fof(f150, plain, ~ (op1(e10, e11) = op1(e13, e11)), inference(cnf_transformation, [], [f5])).
fof(f2197, plain, (spl12_59 | ~ spl12_2), inference(avatar_split_clause, [], [f2156, f493, f735])).
fof(f2156, plain, ((e12 = op1(e10, e11)) | ~ spl12_2), inference(backward_demodulation, [], [f1225, f495])).
fof(f1225, plain, (e12 = op1(e10, op1(e13, e13))), inference(backward_demodulation, [], [f396, f394])).
fof(f394, plain, (e10 = op1(e13, op1(e13, e13))), inference(cnf_transformation, [], [f12])).
fof(f12, plain, ((e12 = op1(op1(e13, op1(e13, e13)), op1(e13, e13))) & (e11 = op1(e13, e13)) & (e10 = op1(e13, op1(e13, e13)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG111+1.p', ax12)).
fof(f396, plain, (e12 = op1(op1(e13, op1(e13, e13)), op1(e13, e13))), inference(cnf_transformation, [], [f12])).
fof(f2184, plain, (~ spl12_33 | ~ spl12_17), inference(avatar_split_clause, [], [f2183, f557, f625])).
fof(f625, plain, (spl12_33 <=> (e10 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl12_33])])).
fof(f2183, plain, (~ (e10 = op1(e11, e13)) | ~ spl12_17), inference(forward_demodulation, [], [f163, f559])).
fof(f163, plain, ~ (op1(e11, e13) = op1(e12, e13)), inference(cnf_transformation, [], [f5])).
fof(f2160, plain, (~ spl12_2 | ~ spl12_58), inference(avatar_contradiction_clause, [], [f2159])).
fof(f2159, plain, ($false | (~ spl12_2 | ~ spl12_58)), inference(subsumption_resolution, [], [f2158, f241])).
fof(f241, plain, ~ (e11 = e12), inference(cnf_transformation, [], [f7])).
fof(f2158, plain, ((e11 = e12) | (~ spl12_2 | ~ spl12_58)), inference(forward_demodulation, [], [f2156, f733])).
fof(f733, plain, ((e11 = op1(e10, e11)) | ~ spl12_58), inference(avatar_component_clause, [], [f731])).
fof(f731, plain, (spl12_58 <=> (e11 = op1(e10, e11))), introduced(avatar_definition, [new_symbols(naming, [spl12_58])])).
fof(f2157, plain, (spl12_9 | ~ spl12_2), inference(avatar_split_clause, [], [f2153, f493, f523])).
fof(f2153, plain, ((e10 = op1(e13, e11)) | ~ spl12_2), inference(backward_demodulation, [], [f394, f495])).
fof(f2146, plain, (~ spl12_7 | ~ spl12_8), inference(avatar_contradiction_clause, [], [f2145])).
fof(f2145, plain, ($false | (~ spl12_7 | ~ spl12_8)), inference(subsumption_resolution, [], [f2144, f243])).
fof(f243, plain, ~ (e12 = e13), inference(cnf_transformation, [], [f7])).
fof(f2144, plain, ((e12 = e13) | (~ spl12_7 | ~ spl12_8)), inference(backward_demodulation, [], [f520, f516])).
fof(f520, plain, ((e13 = op1(e13, e12)) | ~ spl12_8), inference(avatar_component_clause, [], [f518])).
fof(f518, plain, (spl12_8 <=> (e13 = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl12_8])])).
fof(f2082, plain, (~ spl12_9 | ~ spl12_41), inference(avatar_split_clause, [], [f2074, f659, f523])).
fof(f659, plain, (spl12_41 <=> (e10 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl12_41])])).
fof(f2074, plain, (~ (e10 = op1(e13, e11)) | ~ spl12_41), inference(backward_demodulation, [], [f152, f661])).
fof(f661, plain, ((e10 = op1(e11, e11)) | ~ spl12_41), inference(avatar_component_clause, [], [f659])).
fof(f152, plain, ~ (op1(e11, e11) = op1(e13, e11)), inference(cnf_transformation, [], [f5])).
fof(f2050, plain, (~ spl12_21 | ~ spl12_53), inference(avatar_split_clause, [], [f2046, f710, f574])).
fof(f574, plain, (spl12_21 <=> (e10 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl12_21])])).
fof(f2046, plain, (~ (e10 = op1(e12, e12)) | ~ spl12_53), inference(backward_demodulation, [], [f155, f712])).
fof(f155, plain, ~ (op1(e10, e12) = op1(e12, e12)), inference(cnf_transformation, [], [f5])).
fof(f2049, plain, (~ spl12_37 | ~ spl12_53), inference(avatar_split_clause, [], [f2045, f710, f642])).
fof(f642, plain, (spl12_37 <=> (e10 = op1(e11, e12))), introduced(avatar_definition, [new_symbols(naming, [spl12_37])])).
fof(f2045, plain, (~ (e10 = op1(e11, e12)) | ~ spl12_53), inference(backward_demodulation, [], [f154, f712])).
fof(f154, plain, ~ (op1(e10, e12) = op1(e11, e12)), inference(cnf_transformation, [], [f5])).
fof(f2014, plain, (spl12_130 | ~ spl12_66), inference(avatar_split_clause, [], [f2013, f797, f1276])).
fof(f797, plain, (spl12_66 <=> (e21 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl12_66])])).
fof(f2013, plain, ((e21 = h4(e11)) | ~ spl12_66), inference(backward_demodulation, [], [f414, f799])).
fof(f799, plain, ((e21 = op2(e23, e23)) | ~ spl12_66), inference(avatar_component_clause, [], [f797])).
fof(f2007, plain, (~ spl12_71 | ~ spl12_72), inference(avatar_contradiction_clause, [], [f2006])).
fof(f2006, plain, ($false | (~ spl12_71 | ~ spl12_72)), inference(subsumption_resolution, [], [f2005, f249])).
fof(f2005, plain, ((e22 = e23) | (~ spl12_71 | ~ spl12_72)), inference(backward_demodulation, [], [f824, f820])).
fof(f824, plain, ((e23 = op2(e23, e22)) | ~ spl12_72), inference(avatar_component_clause, [], [f822])).
fof(f822, plain, (spl12_72 <=> (e23 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl12_72])])).
fof(f1981, plain, (spl12_139 | ~ spl12_85), inference(avatar_split_clause, [], [f1980, f878, f1326])).
fof(f1326, plain, (spl12_139 <=> (e20 = h3(e11))), introduced(avatar_definition, [new_symbols(naming, [spl12_139])])).
fof(f878, plain, (spl12_85 <=> (e20 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl12_85])])).
fof(f1980, plain, ((e20 = h3(e11)) | ~ spl12_85), inference(backward_demodulation, [], [f410, f880])).
fof(f880, plain, ((e20 = op2(e22, e22)) | ~ spl12_85), inference(avatar_component_clause, [], [f878])).
fof(f1919, plain, (~ spl12_84 | ~ spl12_116), inference(avatar_split_clause, [], [f1915, f1009, f873])).
fof(f873, plain, (spl12_84 <=> (e23 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl12_84])])).
fof(f1915, plain, (~ (e23 = op2(e22, e23)) | ~ spl12_116), inference(backward_demodulation, [], [f209, f1011])).
fof(f209, plain, ~ (op2(e20, e23) = op2(e22, e23)), inference(cnf_transformation, [], [f6])).
fof(f1908, plain, (~ spl12_139 | ~ spl12_117), inference(avatar_split_clause, [], [f1903, f1014, f1326])).
fof(f1903, plain, (~ (e20 = h3(e11)) | ~ spl12_117), inference(backward_demodulation, [], [f1246, f1016])).
fof(f1246, plain, ~ (op2(e20, e22) = h3(e11)), inference(backward_demodulation, [], [f203, f410])).
fof(f203, plain, ~ (op2(e20, e22) = op2(e22, e22)), inference(cnf_transformation, [], [f6])).
fof(f1891, plain, (spl12_155 | ~ spl12_126), inference(avatar_split_clause, [], [f1890, f1052, f1407])).
fof(f1052, plain, (spl12_126 <=> (op2(e20, e20) = e21)), introduced(avatar_definition, [new_symbols(naming, [spl12_126])])).
fof(f1890, plain, ((e21 = h1(e11)) | ~ spl12_126), inference(backward_demodulation, [], [f402, f1054])).
fof(f1054, plain, ((op2(e20, e20) = e21) | ~ spl12_126), inference(avatar_component_clause, [], [f1052])).
fof(f1509, plain, (~ spl12_157 | spl12_131 | spl12_129 | ~ spl12_158 | ~ spl12_159 | ~ spl12_160 | ~ spl12_161 | ~ spl12_162 | ~ spl12_163 | ~ spl12_164 | ~ spl12_165 | ~ spl12_166 | ~ spl12_167 | ~ spl12_168 | ~ spl12_169 | ~ spl12_170 | ~ spl12_171 | ~ spl12_172), inference(avatar_split_clause, [], [f1444, f1506, f1502, f1498, f1494, f1490, f1486, f1482, f1478, f1474, f1470, f1466, f1462, f1458, f1454, f1450, f1272, f1284, f1446])).
fof(f1284, plain, (spl12_131 <=> sP9), introduced(avatar_definition, [new_symbols(naming, [spl12_131])])).
fof(f1272, plain, (spl12_129 <=> sP10), introduced(avatar_definition, [new_symbols(naming, [spl12_129])])).
fof(f1444, plain, (~ (h1(e11) = h4(op1(e10, e10))) | ~ (e22 = h4(op1(e10, e11))) | ~ (op2(e20, e22) = h4(op1(e10, e12))) | ~ (op2(e20, e23) = h4(op1(e10, e13))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), e20)) | ~ (h4(op1(e11, e12)) = op2(h4(e11), e22)) | ~ (h4(op1(e11, e13)) = op2(h4(e11), e23)) | ~ (op2(e22, e20) = h4(op1(e12, e10))) | ~ (h4(op1(e12, e11)) = op2(e22, h4(e11))) | ~ (h3(e11) = h4(op1(e12, e12))) | ~ (op2(e22, e23) = h4(op1(e12, e13))) | ~ (op2(e23, e20) = h4(op1(e13, e10))) | ~ (e20 = h4(op1(e13, e11))) | ~ (op2(e23, e22) = h4(op1(e13, e12))) | ~ (h4(e11) = h4(op1(e13, e13))) | sP10 | sP9 | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11)))), inference(forward_demodulation, [], [f1443, f402])).
fof(f1443, plain, (~ (op2(e20, e20) = h4(op1(e10, e10))) | ~ (e22 = h4(op1(e10, e11))) | ~ (op2(e20, e22) = h4(op1(e10, e12))) | ~ (op2(e20, e23) = h4(op1(e10, e13))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), e20)) | ~ (h4(op1(e11, e12)) = op2(h4(e11), e22)) | ~ (h4(op1(e11, e13)) = op2(h4(e11), e23)) | ~ (op2(e22, e20) = h4(op1(e12, e10))) | ~ (h4(op1(e12, e11)) = op2(e22, h4(e11))) | ~ (h3(e11) = h4(op1(e12, e12))) | ~ (op2(e22, e23) = h4(op1(e12, e13))) | ~ (op2(e23, e20) = h4(op1(e13, e10))) | ~ (e20 = h4(op1(e13, e11))) | ~ (op2(e23, e22) = h4(op1(e13, e12))) | ~ (h4(e11) = h4(op1(e13, e13))) | sP10 | sP9 | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11)))), inference(forward_demodulation, [], [f1442, f1266])).
fof(f1442, plain, (~ (e22 = h4(op1(e10, e11))) | ~ (op2(e20, e22) = h4(op1(e10, e12))) | ~ (op2(e20, e23) = h4(op1(e10, e13))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), e20)) | ~ (h4(op1(e11, e12)) = op2(h4(e11), e22)) | ~ (h4(op1(e11, e13)) = op2(h4(e11), e23)) | ~ (op2(e22, e20) = h4(op1(e12, e10))) | ~ (h4(op1(e12, e11)) = op2(e22, h4(e11))) | ~ (h3(e11) = h4(op1(e12, e12))) | ~ (op2(e22, e23) = h4(op1(e12, e13))) | ~ (op2(e23, e20) = h4(op1(e13, e10))) | ~ (e20 = h4(op1(e13, e11))) | ~ (op2(e23, e22) = h4(op1(e13, e12))) | ~ (h4(e11) = h4(op1(e13, e13))) | sP10 | sP9 | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f1441, f1264])).
fof(f1441, plain, (~ (h4(op1(e10, e11)) = op2(e20, h4(e11))) | ~ (op2(e20, e22) = h4(op1(e10, e12))) | ~ (op2(e20, e23) = h4(op1(e10, e13))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), e20)) | ~ (h4(op1(e11, e12)) = op2(h4(e11), e22)) | ~ (h4(op1(e11, e13)) = op2(h4(e11), e23)) | ~ (op2(e22, e20) = h4(op1(e12, e10))) | ~ (h4(op1(e12, e11)) = op2(e22, h4(e11))) | ~ (h3(e11) = h4(op1(e12, e12))) | ~ (op2(e22, e23) = h4(op1(e12, e13))) | ~ (op2(e23, e20) = h4(op1(e13, e10))) | ~ (e20 = h4(op1(e13, e11))) | ~ (op2(e23, e22) = h4(op1(e13, e12))) | ~ (h4(e11) = h4(op1(e13, e13))) | sP10 | sP9 | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f1440, f1266])).
fof(f1440, plain, (~ (op2(e20, e22) = h4(op1(e10, e12))) | ~ (op2(e20, e23) = h4(op1(e10, e13))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), e20)) | ~ (h4(op1(e11, e12)) = op2(h4(e11), e22)) | ~ (h4(op1(e11, e13)) = op2(h4(e11), e23)) | ~ (op2(e22, e20) = h4(op1(e12, e10))) | ~ (h4(op1(e12, e11)) = op2(e22, h4(e11))) | ~ (h3(e11) = h4(op1(e12, e12))) | ~ (op2(e22, e23) = h4(op1(e12, e13))) | ~ (op2(e23, e20) = h4(op1(e13, e10))) | ~ (e20 = h4(op1(e13, e11))) | ~ (op2(e23, e22) = h4(op1(e13, e12))) | ~ (h4(e11) = h4(op1(e13, e13))) | sP10 | sP9 | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f1439, f1266])).
fof(f1439, plain, (~ (h4(op1(e10, e12)) = op2(h4(e10), e22)) | ~ (op2(e20, e23) = h4(op1(e10, e13))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), e20)) | ~ (h4(op1(e11, e12)) = op2(h4(e11), e22)) | ~ (h4(op1(e11, e13)) = op2(h4(e11), e23)) | ~ (op2(e22, e20) = h4(op1(e12, e10))) | ~ (h4(op1(e12, e11)) = op2(e22, h4(e11))) | ~ (h3(e11) = h4(op1(e12, e12))) | ~ (op2(e22, e23) = h4(op1(e12, e13))) | ~ (op2(e23, e20) = h4(op1(e13, e10))) | ~ (e20 = h4(op1(e13, e11))) | ~ (op2(e23, e22) = h4(op1(e13, e12))) | ~ (h4(e11) = h4(op1(e13, e13))) | sP10 | sP9 | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f1438, f1256])).
fof(f1438, plain, (~ (op2(e20, e23) = h4(op1(e10, e13))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), e20)) | ~ (h4(op1(e11, e12)) = op2(h4(e11), e22)) | ~ (h4(op1(e11, e13)) = op2(h4(e11), e23)) | ~ (op2(e22, e20) = h4(op1(e12, e10))) | ~ (h4(op1(e12, e11)) = op2(e22, h4(e11))) | ~ (h3(e11) = h4(op1(e12, e12))) | ~ (op2(e22, e23) = h4(op1(e12, e13))) | ~ (op2(e23, e20) = h4(op1(e13, e10))) | ~ (e20 = h4(op1(e13, e11))) | ~ (op2(e23, e22) = h4(op1(e13, e12))) | ~ (h4(e11) = h4(op1(e13, e13))) | sP10 | sP9 | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f1437, f1266])).
fof(f1437, plain, (~ (h4(op1(e10, e13)) = op2(h4(e10), e23)) | ~ (h4(op1(e11, e10)) = op2(h4(e11), e20)) | ~ (h4(op1(e11, e12)) = op2(h4(e11), e22)) | ~ (h4(op1(e11, e13)) = op2(h4(e11), e23)) | ~ (op2(e22, e20) = h4(op1(e12, e10))) | ~ (h4(op1(e12, e11)) = op2(e22, h4(e11))) | ~ (h3(e11) = h4(op1(e12, e12))) | ~ (op2(e22, e23) = h4(op1(e12, e13))) | ~ (op2(e23, e20) = h4(op1(e13, e10))) | ~ (e20 = h4(op1(e13, e11))) | ~ (op2(e23, e22) = h4(op1(e13, e12))) | ~ (h4(e11) = h4(op1(e13, e13))) | sP10 | sP9 | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f1436, f412])).
fof(f1436, plain, (~ (h4(op1(e11, e10)) = op2(h4(e11), e20)) | ~ (h4(op1(e11, e12)) = op2(h4(e11), e22)) | ~ (h4(op1(e11, e13)) = op2(h4(e11), e23)) | ~ (op2(e22, e20) = h4(op1(e12, e10))) | ~ (h4(op1(e12, e11)) = op2(e22, h4(e11))) | ~ (h3(e11) = h4(op1(e12, e12))) | ~ (op2(e22, e23) = h4(op1(e12, e13))) | ~ (op2(e23, e20) = h4(op1(e13, e10))) | ~ (e20 = h4(op1(e13, e11))) | ~ (op2(e23, e22) = h4(op1(e13, e12))) | ~ (h4(e11) = h4(op1(e13, e13))) | sP10 | sP9 | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f1435, f1266])).
fof(f1435, plain, (~ (h4(op1(e11, e12)) = op2(h4(e11), e22)) | ~ (h4(op1(e11, e13)) = op2(h4(e11), e23)) | ~ (op2(e22, e20) = h4(op1(e12, e10))) | ~ (h4(op1(e12, e11)) = op2(e22, h4(e11))) | ~ (h3(e11) = h4(op1(e12, e12))) | ~ (op2(e22, e23) = h4(op1(e12, e13))) | ~ (op2(e23, e20) = h4(op1(e13, e10))) | ~ (e20 = h4(op1(e13, e11))) | ~ (op2(e23, e22) = h4(op1(e13, e12))) | ~ (h4(e11) = h4(op1(e13, e13))) | sP10 | sP9 | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f1434, f1256])).
fof(f1434, plain, (~ (h4(op1(e11, e13)) = op2(h4(e11), e23)) | ~ (op2(e22, e20) = h4(op1(e12, e10))) | ~ (h4(op1(e12, e11)) = op2(e22, h4(e11))) | ~ (h3(e11) = h4(op1(e12, e12))) | ~ (op2(e22, e23) = h4(op1(e12, e13))) | ~ (op2(e23, e20) = h4(op1(e13, e10))) | ~ (e20 = h4(op1(e13, e11))) | ~ (op2(e23, e22) = h4(op1(e13, e12))) | ~ (h4(e11) = h4(op1(e13, e13))) | sP10 | sP9 | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f1433, f412])).
fof(f1433, plain, (~ (op2(e22, e20) = h4(op1(e12, e10))) | ~ (h4(op1(e12, e11)) = op2(e22, h4(e11))) | ~ (h3(e11) = h4(op1(e12, e12))) | ~ (op2(e22, e23) = h4(op1(e12, e13))) | ~ (op2(e23, e20) = h4(op1(e13, e10))) | ~ (e20 = h4(op1(e13, e11))) | ~ (op2(e23, e22) = h4(op1(e13, e12))) | ~ (h4(e11) = h4(op1(e13, e13))) | sP10 | sP9 | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f1432, f1256])).
fof(f1432, plain, (~ (h4(op1(e12, e10)) = op2(h4(e12), e20)) | ~ (h4(op1(e12, e11)) = op2(e22, h4(e11))) | ~ (h3(e11) = h4(op1(e12, e12))) | ~ (op2(e22, e23) = h4(op1(e12, e13))) | ~ (op2(e23, e20) = h4(op1(e13, e10))) | ~ (e20 = h4(op1(e13, e11))) | ~ (op2(e23, e22) = h4(op1(e13, e12))) | ~ (h4(e11) = h4(op1(e13, e13))) | sP10 | sP9 | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f1431, f1266])).
fof(f1431, plain, (~ (h4(op1(e12, e11)) = op2(e22, h4(e11))) | ~ (h3(e11) = h4(op1(e12, e12))) | ~ (op2(e22, e23) = h4(op1(e12, e13))) | ~ (op2(e23, e20) = h4(op1(e13, e10))) | ~ (e20 = h4(op1(e13, e11))) | ~ (op2(e23, e22) = h4(op1(e13, e12))) | ~ (h4(e11) = h4(op1(e13, e13))) | sP10 | sP9 | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f1430, f1256])).
fof(f1430, plain, (~ (h3(e11) = h4(op1(e12, e12))) | ~ (op2(e22, e23) = h4(op1(e12, e13))) | ~ (op2(e23, e20) = h4(op1(e13, e10))) | ~ (e20 = h4(op1(e13, e11))) | ~ (op2(e23, e22) = h4(op1(e13, e12))) | ~ (h4(e11) = h4(op1(e13, e13))) | sP10 | sP9 | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f1429, f410])).
fof(f1429, plain, (~ (op2(e22, e22) = h4(op1(e12, e12))) | ~ (op2(e22, e23) = h4(op1(e12, e13))) | ~ (op2(e23, e20) = h4(op1(e13, e10))) | ~ (e20 = h4(op1(e13, e11))) | ~ (op2(e23, e22) = h4(op1(e13, e12))) | ~ (h4(e11) = h4(op1(e13, e13))) | sP10 | sP9 | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f1428, f1256])).
fof(f1428, plain, (~ (op2(e22, e23) = h4(op1(e12, e13))) | ~ (op2(e23, e20) = h4(op1(e13, e10))) | ~ (e20 = h4(op1(e13, e11))) | ~ (op2(e23, e22) = h4(op1(e13, e12))) | ~ (h4(e11) = h4(op1(e13, e13))) | sP10 | sP9 | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f1427, f1256])).
fof(f1427, plain, (~ (h4(op1(e12, e13)) = op2(h4(e12), e23)) | ~ (op2(e23, e20) = h4(op1(e13, e10))) | ~ (e20 = h4(op1(e13, e11))) | ~ (op2(e23, e22) = h4(op1(e13, e12))) | ~ (h4(e11) = h4(op1(e13, e13))) | sP10 | sP9 | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f1426, f412])).
fof(f1426, plain, (~ (op2(e23, e20) = h4(op1(e13, e10))) | ~ (e20 = h4(op1(e13, e11))) | ~ (op2(e23, e22) = h4(op1(e13, e12))) | ~ (h4(e11) = h4(op1(e13, e13))) | sP10 | sP9 | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f1425, f412])).
fof(f1425, plain, (~ (h4(op1(e13, e10)) = op2(h4(e13), e20)) | ~ (e20 = h4(op1(e13, e11))) | ~ (op2(e23, e22) = h4(op1(e13, e12))) | ~ (h4(e11) = h4(op1(e13, e13))) | sP10 | sP9 | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f1424, f1266])).
fof(f1424, plain, (~ (e20 = h4(op1(e13, e11))) | ~ (op2(e23, e22) = h4(op1(e13, e12))) | ~ (h4(e11) = h4(op1(e13, e13))) | sP10 | sP9 | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f1423, f1263])).
fof(f1423, plain, (~ (h4(op1(e13, e11)) = op2(e23, h4(e11))) | ~ (op2(e23, e22) = h4(op1(e13, e12))) | ~ (h4(e11) = h4(op1(e13, e13))) | sP10 | sP9 | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f1422, f412])).
fof(f1422, plain, (~ (op2(e23, e22) = h4(op1(e13, e12))) | ~ (h4(e11) = h4(op1(e13, e13))) | sP10 | sP9 | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f1421, f412])).
fof(f1421, plain, (~ (h4(op1(e13, e12)) = op2(h4(e13), e22)) | ~ (h4(e11) = h4(op1(e13, e13))) | sP10 | sP9 | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f1420, f1256])).
fof(f1420, plain, (~ (h4(e11) = h4(op1(e13, e13))) | sP10 | sP9 | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f1419, f414])).
fof(f1419, plain, (~ (op2(e23, e23) = h4(op1(e13, e13))) | sP10 | sP9 | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f1418, f412])).
fof(f1418, plain, (sP10 | sP9 | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(subsumption_resolution, [], [f1417, f1268])).
fof(f1268, plain, ~ sP11, inference(subsumption_resolution, [], [f418, f1256])).
fof(f418, plain, (~ (e22 = h4(e12)) | ~ sP11), inference(cnf_transformation, [], [f34])).
fof(f34, plain, ((~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10))) | ~ sP11), inference(nnf_transformation, [], [f32])).
fof(f32, plain, ((~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10))) | ~ sP11), inference(usedef, [], [e32])).
fof(e32, plain, (sP11 <=> (~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f1417, plain, (sP11 | sP10 | sP9 | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(subsumption_resolution, [], [f479, f412])).
fof(f479, plain, (~ (e23 = h4(e13)) | sP11 | sP10 | sP9 | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(cnf_transformation, [], [f33])).
fof(f33, plain, (((~ (e23 = h4(e13)) & ~ (e23 = h4(e12)) & ~ (e23 = h4(e11)) & ~ (e23 = h4(e10))) | sP11 | sP10 | sP9 | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) & ((~ (e23 = h3(e13)) & ~ (e23 = h3(e12)) & ~ (e23 = h3(e11)) & ~ (e23 = h3(e10))) | sP8 | sP7 | sP6 | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) & ((~ (e23 = h2(e13)) & ~ (e23 = h2(e12)) & ~ (e23 = h2(e11)) & ~ (e23 = h2(e10))) | sP5 | sP4 | sP3 | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) & ((~ (e23 = h1(e13)) & ~ (e23 = h1(e12)) & ~ (e23 = h1(e11)) & ~ (e23 = h1(e10))) | sP2 | sP1 | sP0 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(definition_folding, [], [f20, e32, e31, e30, e29, e28, e27, e26, e25, e24, e23, e22, e21])).
fof(f21, plain, ((~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10))) | ~ sP0), inference(usedef, [], [e21])).
fof(e21, plain, (sP0 <=> (~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f22, plain, ((~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10))) | ~ sP1), inference(usedef, [], [e22])).
fof(e22, plain, (sP1 <=> (~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f23, plain, ((~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10))) | ~ sP2), inference(usedef, [], [e23])).
fof(e23, plain, (sP2 <=> (~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f24, plain, ((~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10))) | ~ sP3), inference(usedef, [], [e24])).
fof(e24, plain, (sP3 <=> (~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f25, plain, ((~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10))) | ~ sP4), inference(usedef, [], [e25])).
fof(e25, plain, (sP4 <=> (~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f26, plain, ((~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10))) | ~ sP5), inference(usedef, [], [e26])).
fof(e26, plain, (sP5 <=> (~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f27, plain, ((~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10))) | ~ sP6), inference(usedef, [], [e27])).
fof(e27, plain, (sP6 <=> (~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f28, plain, ((~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10))) | ~ sP7), inference(usedef, [], [e28])).
fof(e28, plain, (sP7 <=> (~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f29, plain, ((~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10))) | ~ sP8), inference(usedef, [], [e29])).
fof(e29, plain, (sP8 <=> (~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f30, plain, ((~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10))) | ~ sP9), inference(usedef, [], [e30])).
fof(e30, plain, (sP9 <=> (~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f31, plain, ((~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10))) | ~ sP10), inference(usedef, [], [e31])).
fof(e31, plain, (sP10 <=> (~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f20, plain, (((~ (e23 = h4(e13)) & ~ (e23 = h4(e12)) & ~ (e23 = h4(e11)) & ~ (e23 = h4(e10))) | (~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10))) | (~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10))) | (~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10))) | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) & ((~ (e23 = h3(e13)) & ~ (e23 = h3(e12)) & ~ (e23 = h3(e11)) & ~ (e23 = h3(e10))) | (~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10))) | (~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10))) | (~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10))) | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) & ((~ (e23 = h2(e13)) & ~ (e23 = h2(e12)) & ~ (e23 = h2(e11)) & ~ (e23 = h2(e10))) | (~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10))) | (~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10))) | (~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) & ((~ (e23 = h1(e13)) & ~ (e23 = h1(e12)) & ~ (e23 = h1(e11)) & ~ (e23 = h1(e10))) | (~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10))) | (~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10))) | (~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10))) | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(ennf_transformation, [], [f19])).
fof(f19, plain, ~ ((((e23 = h4(e13)) | (e23 = h4(e12)) | (e23 = h4(e11)) | (e23 = h4(e10))) & ((e22 = h4(e13)) | (e22 = h4(e12)) | (e22 = h4(e11)) | (e22 = h4(e10))) & ((e21 = h4(e13)) | (e21 = h4(e12)) | (e21 = h4(e11)) | (e21 = h4(e10))) & ((e20 = h4(e13)) | (e20 = h4(e12)) | (e20 = h4(e11)) | (e20 = h4(e10))) & (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) & (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) & (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) & (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) & (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) & (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) & (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) & (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) & (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) & (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) & (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) & (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) & (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) & (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) & (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) & (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) | (((e23 = h3(e13)) | (e23 = h3(e12)) | (e23 = h3(e11)) | (e23 = h3(e10))) & ((e22 = h3(e13)) | (e22 = h3(e12)) | (e22 = h3(e11)) | (e22 = h3(e10))) & ((e21 = h3(e13)) | (e21 = h3(e12)) | (e21 = h3(e11)) | (e21 = h3(e10))) & ((e20 = h3(e13)) | (e20 = h3(e12)) | (e20 = h3(e11)) | (e20 = h3(e10))) & (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) & (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) & (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) & (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) & (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) & (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) & (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) & (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) & (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) & (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) & (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) & (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) & (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) & (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) & (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) & (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) | (((e23 = h2(e13)) | (e23 = h2(e12)) | (e23 = h2(e11)) | (e23 = h2(e10))) & ((e22 = h2(e13)) | (e22 = h2(e12)) | (e22 = h2(e11)) | (e22 = h2(e10))) & ((e21 = h2(e13)) | (e21 = h2(e12)) | (e21 = h2(e11)) | (e21 = h2(e10))) & ((e20 = h2(e13)) | (e20 = h2(e12)) | (e20 = h2(e11)) | (e20 = h2(e10))) & (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) & (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) & (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) & (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) & (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) & (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) & (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) & (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) & (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) & (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) & (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) & (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) & (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) & (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) & (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) & (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) | (((e23 = h1(e13)) | (e23 = h1(e12)) | (e23 = h1(e11)) | (e23 = h1(e10))) & ((e22 = h1(e13)) | (e22 = h1(e12)) | (e22 = h1(e11)) | (e22 = h1(e10))) & ((e21 = h1(e13)) | (e21 = h1(e12)) | (e21 = h1(e11)) | (e21 = h1(e10))) & ((e20 = h1(e13)) | (e20 = h1(e12)) | (e20 = h1(e11)) | (e20 = h1(e10))) & (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) & (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) & (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) & (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) & (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) & (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) & (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) & (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) & (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) & (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) & (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) & (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) & (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) & (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) & (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) & (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(negated_conjecture, [], [f18])).
fof(f18, plain, ~ ((((e23 = h4(e13)) | (e23 = h4(e12)) | (e23 = h4(e11)) | (e23 = h4(e10))) & ((e22 = h4(e13)) | (e22 = h4(e12)) | (e22 = h4(e11)) | (e22 = h4(e10))) & ((e21 = h4(e13)) | (e21 = h4(e12)) | (e21 = h4(e11)) | (e21 = h4(e10))) & ((e20 = h4(e13)) | (e20 = h4(e12)) | (e20 = h4(e11)) | (e20 = h4(e10))) & (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) & (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) & (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) & (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) & (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) & (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) & (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) & (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) & (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) & (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) & (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) & (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) & (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) & (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) & (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) & (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) | (((e23 = h3(e13)) | (e23 = h3(e12)) | (e23 = h3(e11)) | (e23 = h3(e10))) & ((e22 = h3(e13)) | (e22 = h3(e12)) | (e22 = h3(e11)) | (e22 = h3(e10))) & ((e21 = h3(e13)) | (e21 = h3(e12)) | (e21 = h3(e11)) | (e21 = h3(e10))) & ((e20 = h3(e13)) | (e20 = h3(e12)) | (e20 = h3(e11)) | (e20 = h3(e10))) & (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) & (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) & (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) & (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) & (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) & (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) & (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) & (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) & (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) & (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) & (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) & (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) & (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) & (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) & (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) & (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) | (((e23 = h2(e13)) | (e23 = h2(e12)) | (e23 = h2(e11)) | (e23 = h2(e10))) & ((e22 = h2(e13)) | (e22 = h2(e12)) | (e22 = h2(e11)) | (e22 = h2(e10))) & ((e21 = h2(e13)) | (e21 = h2(e12)) | (e21 = h2(e11)) | (e21 = h2(e10))) & ((e20 = h2(e13)) | (e20 = h2(e12)) | (e20 = h2(e11)) | (e20 = h2(e10))) & (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) & (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) & (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) & (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) & (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) & (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) & (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) & (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) & (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) & (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) & (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) & (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) & (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) & (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) & (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) & (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) | (((e23 = h1(e13)) | (e23 = h1(e12)) | (e23 = h1(e11)) | (e23 = h1(e10))) & ((e22 = h1(e13)) | (e22 = h1(e12)) | (e22 = h1(e11)) | (e22 = h1(e10))) & ((e21 = h1(e13)) | (e21 = h1(e12)) | (e21 = h1(e11)) | (e21 = h1(e10))) & ((e20 = h1(e13)) | (e20 = h1(e12)) | (e20 = h1(e11)) | (e20 = h1(e10))) & (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) & (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) & (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) & (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) & (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) & (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) & (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) & (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) & (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) & (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) & (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) & (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) & (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) & (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) & (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) & (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG111+1.p', co1)).
fof(f1293, plain, ~ spl12_131, inference(avatar_split_clause, [], [f1292, f1284])).
fof(f1292, plain, ~ sP9, inference(subsumption_resolution, [], [f424, f1266])).
fof(f424, plain, (~ (e20 = h4(e10)) | ~ sP9), inference(cnf_transformation, [], [f36])).
fof(f36, plain, ((~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10))) | ~ sP9), inference(nnf_transformation, [], [f30])).
fof(f1279, plain, (~ spl12_129 | ~ spl12_130), inference(avatar_split_clause, [], [f421, f1276, f1272])).
fof(f421, plain, (~ (e21 = h4(e11)) | ~ sP10), inference(cnf_transformation, [], [f35])).
fof(f35, plain, ((~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10))) | ~ sP10), inference(nnf_transformation, [], [f31])).
fof(f1226, plain, spl12_66, inference(avatar_split_clause, [], [f398, f797])).
fof(f398, plain, (e21 = op2(e23, e23)), inference(cnf_transformation, [], [f13])).
fof(f1224, plain, spl12_2, inference(avatar_split_clause, [], [f395, f493])).
fof(f395, plain, (e11 = op1(e13, e13)), inference(cnf_transformation, [], [f12])).
fof(f1223, plain, ~ spl12_125, inference(avatar_split_clause, [], [f480, f1048])).
fof(f1048, plain, (spl12_125 <=> (e20 = op2(e20, e20))), introduced(avatar_definition, [new_symbols(naming, [spl12_125])])).
fof(f480, plain, ~ (e20 = op2(e20, e20)), inference(duplicate_literal_removal, [], [f330])).
fof(f330, plain, (~ (e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20))), inference(cnf_transformation, [], [f11])).
fof(f11, plain, ((~ (e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23))) & (~ (e23 = op2(e22, e22)) | ~ (e23 = op2(e23, e23))) & (~ (e23 = op2(e21, e21)) | ~ (e23 = op2(e23, e23))) & (~ (op2(e20, e20) = e23) | ~ (e23 = op2(e23, e23))) & (~ (e22 = op2(e23, e23)) | ~ (e23 = op2(e23, e22))) & (~ (e22 = op2(e22, e22)) | ~ (e23 = op2(e23, e22))) & (~ (e22 = op2(e21, e21)) | ~ (e23 = op2(e23, e22))) & (~ (op2(e20, e20) = e22) | ~ (e23 = op2(e23, e22))) & (~ (e21 = op2(e23, e23)) | ~ (e23 = op2(e23, e21))) & (~ (e21 = op2(e22, e22)) | ~ (e23 = op2(e23, e21))) & (~ (e21 = op2(e21, e21)) | ~ (e23 = op2(e23, e21))) & (~ (op2(e20, e20) = e21) | ~ (e23 = op2(e23, e21))) & (~ (e20 = op2(e23, e23)) | ~ (e23 = op2(e23, e20))) & (~ (e20 = op2(e22, e22)) | ~ (e23 = op2(e23, e20))) & (~ (e20 = op2(e21, e21)) | ~ (e23 = op2(e23, e20))) & (~ (e20 = op2(e20, e20)) | ~ (e23 = op2(e23, e20))) & (~ (e23 = op2(e23, e23)) | ~ (e22 = op2(e22, e23))) & (~ (e23 = op2(e22, e22)) | ~ (e22 = op2(e22, e23))) & (~ (e23 = op2(e21, e21)) | ~ (e22 = op2(e22, e23))) & (~ (op2(e20, e20) = e23) | ~ (e22 = op2(e22, e23))) & (~ (e22 = op2(e23, e23)) | ~ (e22 = op2(e22, e22))) & (~ (e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & (~ (e22 = op2(e21, e21)) | ~ (e22 = op2(e22, e22))) & (~ (op2(e20, e20) = e22) | ~ (e22 = op2(e22, e22))) & (~ (e21 = op2(e23, e23)) | ~ (e22 = op2(e22, e21))) & (~ (e21 = op2(e22, e22)) | ~ (e22 = op2(e22, e21))) & (~ (e21 = op2(e21, e21)) | ~ (e22 = op2(e22, e21))) & (~ (op2(e20, e20) = e21) | ~ (e22 = op2(e22, e21))) & (~ (e20 = op2(e23, e23)) | ~ (e22 = op2(e22, e20))) & (~ (e20 = op2(e22, e22)) | ~ (e22 = op2(e22, e20))) & (~ (e20 = op2(e21, e21)) | ~ (e22 = op2(e22, e20))) & (~ (e20 = op2(e20, e20)) | ~ (e22 = op2(e22, e20))) & (~ (e23 = op2(e23, e23)) | ~ (e21 = op2(e21, e23))) & (~ (e23 = op2(e22, e22)) | ~ (e21 = op2(e21, e23))) & (~ (e23 = op2(e21, e21)) | ~ (e21 = op2(e21, e23))) & (~ (op2(e20, e20) = e23) | ~ (e21 = op2(e21, e23))) & (~ (e22 = op2(e23, e23)) | ~ (e21 = op2(e21, e22))) & (~ (e22 = op2(e22, e22)) | ~ (e21 = op2(e21, e22))) & (~ (e22 = op2(e21, e21)) | ~ (e21 = op2(e21, e22))) & (~ (op2(e20, e20) = e22) | ~ (e21 = op2(e21, e22))) & (~ (e21 = op2(e23, e23)) | ~ (e21 = op2(e21, e21))) & (~ (e21 = op2(e22, e22)) | ~ (e21 = op2(e21, e21))) & (~ (e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & (~ (op2(e20, e20) = e21) | ~ (e21 = op2(e21, e21))) & (~ (e20 = op2(e23, e23)) | ~ (e21 = op2(e21, e20))) & (~ (e20 = op2(e22, e22)) | ~ (e21 = op2(e21, e20))) & (~ (e20 = op2(e21, e21)) | ~ (e21 = op2(e21, e20))) & (~ (e20 = op2(e20, e20)) | ~ (e21 = op2(e21, e20))) & (~ (e23 = op2(e23, e23)) | ~ (e20 = op2(e20, e23))) & (~ (e23 = op2(e22, e22)) | ~ (e20 = op2(e20, e23))) & (~ (e23 = op2(e21, e21)) | ~ (e20 = op2(e20, e23))) & (~ (op2(e20, e20) = e23) | ~ (e20 = op2(e20, e23))) & (~ (e22 = op2(e23, e23)) | ~ (e20 = op2(e20, e22))) & (~ (e22 = op2(e22, e22)) | ~ (e20 = op2(e20, e22))) & (~ (e22 = op2(e21, e21)) | ~ (e20 = op2(e20, e22))) & (~ (op2(e20, e20) = e22) | ~ (e20 = op2(e20, e22))) & (~ (e21 = op2(e23, e23)) | ~ (e20 = op2(e20, e21))) & (~ (e21 = op2(e22, e22)) | ~ (e20 = op2(e20, e21))) & (~ (e21 = op2(e21, e21)) | ~ (e20 = op2(e20, e21))) & (~ (op2(e20, e20) = e21) | ~ (e20 = op2(e20, e21))) & (~ (e20 = op2(e23, e23)) | ~ (e20 = op2(e20, e20))) & (~ (e20 = op2(e22, e22)) | ~ (e20 = op2(e20, e20))) & (~ (e20 = op2(e21, e21)) | ~ (e20 = op2(e20, e20))) & (~ (e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG111+1.p', ax11)).
fof(f1216, plain, (~ spl12_121 | ~ spl12_66), inference(avatar_split_clause, [], [f337, f797, f1031])).
fof(f1031, plain, (spl12_121 <=> (e20 = op2(e20, e21))), introduced(avatar_definition, [new_symbols(naming, [spl12_121])])).
fof(f337, plain, (~ (e21 = op2(e23, e23)) | ~ (e20 = op2(e20, e21))), inference(cnf_transformation, [], [f11])).
fof(f1210, plain, (~ spl12_113 | ~ spl12_108), inference(avatar_split_clause, [], [f343, f975, f997])).
fof(f997, plain, (spl12_113 <=> (e20 = op2(e20, e23))), introduced(avatar_definition, [new_symbols(naming, [spl12_113])])).
fof(f343, plain, (~ (e23 = op2(e21, e21)) | ~ (e20 = op2(e20, e23))), inference(cnf_transformation, [], [f11])).
fof(f1202, plain, ~ spl12_106, inference(avatar_split_clause, [], [f481, f967])).
fof(f967, plain, (spl12_106 <=> (e21 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl12_106])])).
fof(f481, plain, ~ (e21 = op2(e21, e21)), inference(duplicate_literal_removal, [], [f351])).
fof(f351, plain, (~ (e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))), inference(cnf_transformation, [], [f11])).
fof(f1184, plain, (~ spl12_91 | ~ spl12_66), inference(avatar_split_clause, [], [f369, f797, f903])).
fof(f903, plain, (spl12_91 <=> (e22 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl12_91])])).
fof(f369, plain, (~ (e21 = op2(e23, e23)) | ~ (e22 = op2(e22, e21))), inference(cnf_transformation, [], [f11])).
fof(f1181, plain, ~ spl12_87, inference(avatar_split_clause, [], [f482, f886])).
fof(f886, plain, (spl12_87 <=> (e22 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl12_87])])).
fof(f482, plain, ~ (e22 = op2(e22, e22)), inference(duplicate_literal_removal, [], [f372])).
fof(f372, plain, (~ (e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))), inference(cnf_transformation, [], [f11])).
fof(f1178, plain, (~ spl12_83 | ~ spl12_108), inference(avatar_split_clause, [], [f375, f975, f869])).
fof(f869, plain, (spl12_83 <=> (e22 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl12_83])])).
fof(f375, plain, (~ (e23 = op2(e21, e21)) | ~ (e22 = op2(e22, e23))), inference(cnf_transformation, [], [f11])).
fof(f1168, plain, (~ spl12_76 | ~ spl12_66), inference(avatar_split_clause, [], [f385, f797, f839])).
fof(f839, plain, (spl12_76 <=> (e23 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl12_76])])).
fof(f385, plain, (~ (e21 = op2(e23, e23)) | ~ (e23 = op2(e23, e21))), inference(cnf_transformation, [], [f11])).
fof(f1160, plain, ~ spl12_68, inference(avatar_split_clause, [], [f483, f805])).
fof(f805, plain, (spl12_68 <=> (e23 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl12_68])])).
fof(f483, plain, ~ (e23 = op2(e23, e23)), inference(duplicate_literal_removal, [], [f393])).
fof(f393, plain, (~ (e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23))), inference(cnf_transformation, [], [f11])).
fof(f1159, plain, ~ spl12_61, inference(avatar_split_clause, [], [f484, f744])).
fof(f744, plain, (spl12_61 <=> (e10 = op1(e10, e10))), introduced(avatar_definition, [new_symbols(naming, [spl12_61])])).
fof(f484, plain, ~ (e10 = op1(e10, e10)), inference(duplicate_literal_removal, [], [f266])).
fof(f266, plain, (~ (e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10))), inference(cnf_transformation, [], [f10])).
fof(f10, plain, ((~ (e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13))) & (~ (e13 = op1(e12, e12)) | ~ (e13 = op1(e13, e13))) & (~ (e13 = op1(e11, e11)) | ~ (e13 = op1(e13, e13))) & (~ (op1(e10, e10) = e13) | ~ (e13 = op1(e13, e13))) & (~ (e12 = op1(e13, e13)) | ~ (e13 = op1(e13, e12))) & (~ (e12 = op1(e12, e12)) | ~ (e13 = op1(e13, e12))) & (~ (e12 = op1(e11, e11)) | ~ (e13 = op1(e13, e12))) & (~ (op1(e10, e10) = e12) | ~ (e13 = op1(e13, e12))) & (~ (e11 = op1(e13, e13)) | ~ (e13 = op1(e13, e11))) & (~ (e11 = op1(e12, e12)) | ~ (e13 = op1(e13, e11))) & (~ (e11 = op1(e11, e11)) | ~ (e13 = op1(e13, e11))) & (~ (op1(e10, e10) = e11) | ~ (e13 = op1(e13, e11))) & (~ (e10 = op1(e13, e13)) | ~ (e13 = op1(e13, e10))) & (~ (e10 = op1(e12, e12)) | ~ (e13 = op1(e13, e10))) & (~ (e10 = op1(e11, e11)) | ~ (e13 = op1(e13, e10))) & (~ (e10 = op1(e10, e10)) | ~ (e13 = op1(e13, e10))) & (~ (e13 = op1(e13, e13)) | ~ (e12 = op1(e12, e13))) & (~ (e13 = op1(e12, e12)) | ~ (e12 = op1(e12, e13))) & (~ (e13 = op1(e11, e11)) | ~ (e12 = op1(e12, e13))) & (~ (op1(e10, e10) = e13) | ~ (e12 = op1(e12, e13))) & (~ (e12 = op1(e13, e13)) | ~ (e12 = op1(e12, e12))) & (~ (e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & (~ (e12 = op1(e11, e11)) | ~ (e12 = op1(e12, e12))) & (~ (op1(e10, e10) = e12) | ~ (e12 = op1(e12, e12))) & (~ (e11 = op1(e13, e13)) | ~ (e12 = op1(e12, e11))) & (~ (e11 = op1(e12, e12)) | ~ (e12 = op1(e12, e11))) & (~ (e11 = op1(e11, e11)) | ~ (e12 = op1(e12, e11))) & (~ (op1(e10, e10) = e11) | ~ (e12 = op1(e12, e11))) & (~ (e10 = op1(e13, e13)) | ~ (e12 = op1(e12, e10))) & (~ (e10 = op1(e12, e12)) | ~ (e12 = op1(e12, e10))) & (~ (e10 = op1(e11, e11)) | ~ (e12 = op1(e12, e10))) & (~ (e10 = op1(e10, e10)) | ~ (e12 = op1(e12, e10))) & (~ (e13 = op1(e13, e13)) | ~ (e11 = op1(e11, e13))) & (~ (e13 = op1(e12, e12)) | ~ (e11 = op1(e11, e13))) & (~ (e13 = op1(e11, e11)) | ~ (e11 = op1(e11, e13))) & (~ (op1(e10, e10) = e13) | ~ (e11 = op1(e11, e13))) & (~ (e12 = op1(e13, e13)) | ~ (e11 = op1(e11, e12))) & (~ (e12 = op1(e12, e12)) | ~ (e11 = op1(e11, e12))) & (~ (e12 = op1(e11, e11)) | ~ (e11 = op1(e11, e12))) & (~ (op1(e10, e10) = e12) | ~ (e11 = op1(e11, e12))) & (~ (e11 = op1(e13, e13)) | ~ (e11 = op1(e11, e11))) & (~ (e11 = op1(e12, e12)) | ~ (e11 = op1(e11, e11))) & (~ (e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & (~ (op1(e10, e10) = e11) | ~ (e11 = op1(e11, e11))) & (~ (e10 = op1(e13, e13)) | ~ (e11 = op1(e11, e10))) & (~ (e10 = op1(e12, e12)) | ~ (e11 = op1(e11, e10))) & (~ (e10 = op1(e11, e11)) | ~ (e11 = op1(e11, e10))) & (~ (e10 = op1(e10, e10)) | ~ (e11 = op1(e11, e10))) & (~ (e13 = op1(e13, e13)) | ~ (e10 = op1(e10, e13))) & (~ (e13 = op1(e12, e12)) | ~ (e10 = op1(e10, e13))) & (~ (e13 = op1(e11, e11)) | ~ (e10 = op1(e10, e13))) & (~ (op1(e10, e10) = e13) | ~ (e10 = op1(e10, e13))) & (~ (e12 = op1(e13, e13)) | ~ (e10 = op1(e10, e12))) & (~ (e12 = op1(e12, e12)) | ~ (e10 = op1(e10, e12))) & (~ (e12 = op1(e11, e11)) | ~ (e10 = op1(e10, e12))) & (~ (op1(e10, e10) = e12) | ~ (e10 = op1(e10, e12))) & (~ (e11 = op1(e13, e13)) | ~ (e10 = op1(e10, e11))) & (~ (e11 = op1(e12, e12)) | ~ (e10 = op1(e10, e11))) & (~ (e11 = op1(e11, e11)) | ~ (e10 = op1(e10, e11))) & (~ (op1(e10, e10) = e11) | ~ (e10 = op1(e10, e11))) & (~ (e10 = op1(e13, e13)) | ~ (e10 = op1(e10, e10))) & (~ (e10 = op1(e12, e12)) | ~ (e10 = op1(e10, e10))) & (~ (e10 = op1(e11, e11)) | ~ (e10 = op1(e10, e10))) & (~ (e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG111+1.p', ax10)).
fof(f1152, plain, (~ spl12_57 | ~ spl12_2), inference(avatar_split_clause, [], [f273, f493, f727])).
fof(f727, plain, (spl12_57 <=> (e10 = op1(e10, e11))), introduced(avatar_definition, [new_symbols(naming, [spl12_57])])).
fof(f273, plain, (~ (e11 = op1(e13, e13)) | ~ (e10 = op1(e10, e11))), inference(cnf_transformation, [], [f10])).
fof(f1151, plain, (~ spl12_53 | ~ spl12_63), inference(avatar_split_clause, [], [f274, f752, f710])).
fof(f752, plain, (spl12_63 <=> (op1(e10, e10) = e12)), introduced(avatar_definition, [new_symbols(naming, [spl12_63])])).
fof(f274, plain, (~ (op1(e10, e10) = e12) | ~ (e10 = op1(e10, e12))), inference(cnf_transformation, [], [f10])).
fof(f1148, plain, (~ spl12_53 | ~ spl12_3), inference(avatar_split_clause, [], [f277, f497, f710])).
fof(f497, plain, (spl12_3 <=> (e12 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl12_3])])).
fof(f277, plain, (~ (e12 = op1(e13, e13)) | ~ (e10 = op1(e10, e12))), inference(cnf_transformation, [], [f10])).
fof(f1146, plain, (~ spl12_49 | ~ spl12_44), inference(avatar_split_clause, [], [f279, f671, f693])).
fof(f693, plain, (spl12_49 <=> (e10 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl12_49])])).
fof(f279, plain, (~ (e13 = op1(e11, e11)) | ~ (e10 = op1(e10, e13))), inference(cnf_transformation, [], [f10])).
fof(f1138, plain, ~ spl12_42, inference(avatar_split_clause, [], [f485, f663])).
fof(f663, plain, (spl12_42 <=> (e11 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl12_42])])).
fof(f485, plain, ~ (e11 = op1(e11, e11)), inference(duplicate_literal_removal, [], [f287])).
fof(f287, plain, (~ (e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))), inference(cnf_transformation, [], [f10])).
fof(f1120, plain, (~ spl12_27 | ~ spl12_2), inference(avatar_split_clause, [], [f305, f493, f599])).
fof(f599, plain, (spl12_27 <=> (e12 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl12_27])])).
fof(f305, plain, (~ (e11 = op1(e13, e13)) | ~ (e12 = op1(e12, e11))), inference(cnf_transformation, [], [f10])).
fof(f1117, plain, ~ spl12_23, inference(avatar_split_clause, [], [f486, f582])).
fof(f582, plain, (spl12_23 <=> (e12 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl12_23])])).
fof(f486, plain, ~ (e12 = op1(e12, e12)), inference(duplicate_literal_removal, [], [f308])).
fof(f308, plain, (~ (e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))), inference(cnf_transformation, [], [f10])).
fof(f1114, plain, (~ spl12_19 | ~ spl12_44), inference(avatar_split_clause, [], [f311, f671, f565])).
fof(f565, plain, (spl12_19 <=> (e12 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl12_19])])).
fof(f311, plain, (~ (e13 = op1(e11, e11)) | ~ (e12 = op1(e12, e13))), inference(cnf_transformation, [], [f10])).
fof(f1104, plain, (~ spl12_12 | ~ spl12_2), inference(avatar_split_clause, [], [f321, f493, f535])).
fof(f535, plain, (spl12_12 <=> (e13 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl12_12])])).
fof(f321, plain, (~ (e11 = op1(e13, e13)) | ~ (e13 = op1(e13, e11))), inference(cnf_transformation, [], [f10])).
fof(f1096, plain, ~ spl12_4, inference(avatar_split_clause, [], [f487, f501])).
fof(f501, plain, (spl12_4 <=> (e13 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl12_4])])).
fof(f487, plain, ~ (e13 = op1(e13, e13)), inference(duplicate_literal_removal, [], [f329])).
fof(f329, plain, (~ (e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13))), inference(cnf_transformation, [], [f10])).
fof(f1095, plain, (spl12_125 | spl12_121 | spl12_117 | spl12_113), inference(avatar_split_clause, [], [f110, f997, f1014, f1031, f1048])).
fof(f110, plain, ((e20 = op2(e20, e23)) | (e20 = op2(e20, e22)) | (e20 = op2(e20, e21)) | (e20 = op2(e20, e20))), inference(cnf_transformation, [], [f4])).
fof(f4, plain, (((e23 = op2(e23, e23)) | (e23 = op2(e22, e23)) | (e23 = op2(e21, e23)) | (e23 = op2(e20, e23))) & ((e23 = op2(e23, e23)) | (e23 = op2(e23, e22)) | (e23 = op2(e23, e21)) | (e23 = op2(e23, e20))) & ((e22 = op2(e23, e23)) | (e22 = op2(e22, e23)) | (e22 = op2(e21, e23)) | (e22 = op2(e20, e23))) & ((e22 = op2(e23, e23)) | (e22 = op2(e23, e22)) | (e22 = op2(e23, e21)) | (e22 = op2(e23, e20))) & ((e21 = op2(e23, e23)) | (e21 = op2(e22, e23)) | (e21 = op2(e21, e23)) | (e21 = op2(e20, e23))) & ((e21 = op2(e23, e23)) | (e21 = op2(e23, e22)) | (e21 = op2(e23, e21)) | (e21 = op2(e23, e20))) & ((e20 = op2(e23, e23)) | (e20 = op2(e22, e23)) | (e20 = op2(e21, e23)) | (e20 = op2(e20, e23))) & ((e20 = op2(e23, e23)) | (e20 = op2(e23, e22)) | (e20 = op2(e23, e21)) | (e20 = op2(e23, e20))) & ((e23 = op2(e23, e22)) | (e23 = op2(e22, e22)) | (e23 = op2(e21, e22)) | (e23 = op2(e20, e22))) & ((e23 = op2(e22, e23)) | (e23 = op2(e22, e22)) | (e23 = op2(e22, e21)) | (e23 = op2(e22, e20))) & ((e22 = op2(e23, e22)) | (e22 = op2(e22, e22)) | (e22 = op2(e21, e22)) | (e22 = op2(e20, e22))) & ((e22 = op2(e22, e23)) | (e22 = op2(e22, e22)) | (e22 = op2(e22, e21)) | (e22 = op2(e22, e20))) & ((e21 = op2(e23, e22)) | (e21 = op2(e22, e22)) | (e21 = op2(e21, e22)) | (e21 = op2(e20, e22))) & ((e21 = op2(e22, e23)) | (e21 = op2(e22, e22)) | (e21 = op2(e22, e21)) | (e21 = op2(e22, e20))) & ((e20 = op2(e23, e22)) | (e20 = op2(e22, e22)) | (e20 = op2(e21, e22)) | (e20 = op2(e20, e22))) & ((e20 = op2(e22, e23)) | (e20 = op2(e22, e22)) | (e20 = op2(e22, e21)) | (e20 = op2(e22, e20))) & ((e23 = op2(e23, e21)) | (e23 = op2(e22, e21)) | (e23 = op2(e21, e21)) | (e23 = op2(e20, e21))) & ((e23 = op2(e21, e23)) | (e23 = op2(e21, e22)) | (e23 = op2(e21, e21)) | (e23 = op2(e21, e20))) & ((e22 = op2(e23, e21)) | (e22 = op2(e22, e21)) | (e22 = op2(e21, e21)) | (e22 = op2(e20, e21))) & ((e22 = op2(e21, e23)) | (e22 = op2(e21, e22)) | (e22 = op2(e21, e21)) | (e22 = op2(e21, e20))) & ((e21 = op2(e23, e21)) | (e21 = op2(e22, e21)) | (e21 = op2(e21, e21)) | (e21 = op2(e20, e21))) & ((e21 = op2(e21, e23)) | (e21 = op2(e21, e22)) | (e21 = op2(e21, e21)) | (e21 = op2(e21, e20))) & ((e20 = op2(e23, e21)) | (e20 = op2(e22, e21)) | (e20 = op2(e21, e21)) | (e20 = op2(e20, e21))) & ((e20 = op2(e21, e23)) | (e20 = op2(e21, e22)) | (e20 = op2(e21, e21)) | (e20 = op2(e21, e20))) & ((e23 = op2(e23, e20)) | (e23 = op2(e22, e20)) | (e23 = op2(e21, e20)) | (op2(e20, e20) = e23)) & ((e23 = op2(e20, e23)) | (e23 = op2(e20, e22)) | (e23 = op2(e20, e21)) | (op2(e20, e20) = e23)) & ((e22 = op2(e23, e20)) | (e22 = op2(e22, e20)) | (e22 = op2(e21, e20)) | (op2(e20, e20) = e22)) & ((e22 = op2(e20, e23)) | (e22 = op2(e20, e22)) | (e22 = op2(e20, e21)) | (op2(e20, e20) = e22)) & ((e21 = op2(e23, e20)) | (e21 = op2(e22, e20)) | (e21 = op2(e21, e20)) | (op2(e20, e20) = e21)) & ((e21 = op2(e20, e23)) | (e21 = op2(e20, e22)) | (e21 = op2(e20, e21)) | (op2(e20, e20) = e21)) & ((e20 = op2(e23, e20)) | (e20 = op2(e22, e20)) | (e20 = op2(e21, e20)) | (e20 = op2(e20, e20))) & ((e20 = op2(e20, e23)) | (e20 = op2(e20, e22)) | (e20 = op2(e20, e21)) | (e20 = op2(e20, e20)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG111+1.p', ax4)).
fof(f1094, plain, (spl12_125 | spl12_109 | spl12_93 | spl12_77), inference(avatar_split_clause, [], [f111, f844, f912, f980, f1048])).
fof(f111, plain, ((e20 = op2(e23, e20)) | (e20 = op2(e22, e20)) | (e20 = op2(e21, e20)) | (e20 = op2(e20, e20))), inference(cnf_transformation, [], [f4])).
fof(f1093, plain, (spl12_126 | spl12_122 | spl12_118 | spl12_114), inference(avatar_split_clause, [], [f112, f1001, f1018, f1035, f1052])).
fof(f112, plain, ((e21 = op2(e20, e23)) | (e21 = op2(e20, e22)) | (e21 = op2(e20, e21)) | (op2(e20, e20) = e21)), inference(cnf_transformation, [], [f4])).
fof(f1085, plain, (spl12_110 | spl12_106 | spl12_102 | spl12_98), inference(avatar_split_clause, [], [f120, f933, f950, f967, f984])).
fof(f120, plain, ((e21 = op2(e21, e23)) | (e21 = op2(e21, e22)) | (e21 = op2(e21, e21)) | (e21 = op2(e21, e20))), inference(cnf_transformation, [], [f4])).
fof(f1084, plain, (spl12_122 | spl12_106 | spl12_90 | spl12_74), inference(avatar_split_clause, [], [f121, f831, f899, f967, f1035])).
fof(f121, plain, ((e21 = op2(e23, e21)) | (e21 = op2(e22, e21)) | (e21 = op2(e21, e21)) | (e21 = op2(e20, e21))), inference(cnf_transformation, [], [f4])).
fof(f1083, plain, (spl12_111 | spl12_107 | spl12_103 | spl12_99), inference(avatar_split_clause, [], [f122, f937, f954, f971, f988])).
fof(f122, plain, ((e22 = op2(e21, e23)) | (e22 = op2(e21, e22)) | (e22 = op2(e21, e21)) | (e22 = op2(e21, e20))), inference(cnf_transformation, [], [f4])).
fof(f1080, plain, (spl12_124 | spl12_108 | spl12_92 | spl12_76), inference(avatar_split_clause, [], [f125, f839, f907, f975, f1043])).
fof(f125, plain, ((e23 = op2(e23, e21)) | (e23 = op2(e22, e21)) | (e23 = op2(e21, e21)) | (e23 = op2(e20, e21))), inference(cnf_transformation, [], [f4])).
fof(f1079, plain, (spl12_93 | spl12_89 | spl12_85 | spl12_81), inference(avatar_split_clause, [], [f126, f861, f878, f895, f912])).
fof(f126, plain, ((e20 = op2(e22, e23)) | (e20 = op2(e22, e22)) | (e20 = op2(e22, e21)) | (e20 = op2(e22, e20))), inference(cnf_transformation, [], [f4])).
fof(f1075, plain, (spl12_95 | spl12_91 | spl12_87 | spl12_83), inference(avatar_split_clause, [], [f130, f869, f886, f903, f920])).
fof(f130, plain, ((e22 = op2(e22, e23)) | (e22 = op2(e22, e22)) | (e22 = op2(e22, e21)) | (e22 = op2(e22, e20))), inference(cnf_transformation, [], [f4])).
fof(f1074, plain, (spl12_119 | spl12_103 | spl12_87 | spl12_71), inference(avatar_split_clause, [], [f131, f818, f886, f954, f1022])).
fof(f131, plain, ((e22 = op2(e23, e22)) | (e22 = op2(e22, e22)) | (e22 = op2(e21, e22)) | (e22 = op2(e20, e22))), inference(cnf_transformation, [], [f4])).
fof(f1073, plain, (spl12_96 | spl12_92 | spl12_88 | spl12_84), inference(avatar_split_clause, [], [f132, f873, f890, f907, f924])).
fof(f132, plain, ((e23 = op2(e22, e23)) | (e23 = op2(e22, e22)) | (e23 = op2(e22, e21)) | (e23 = op2(e22, e20))), inference(cnf_transformation, [], [f4])).
fof(f1065, plain, (spl12_80 | spl12_76 | spl12_72 | spl12_68), inference(avatar_split_clause, [], [f140, f805, f822, f839, f856])).
fof(f140, plain, ((e23 = op2(e23, e23)) | (e23 = op2(e23, e22)) | (e23 = op2(e23, e21)) | (e23 = op2(e23, e20))), inference(cnf_transformation, [], [f4])).
fof(f1012, plain, (spl12_113 | spl12_114 | spl12_115 | spl12_116), inference(avatar_split_clause, [], [f97, f1009, f1005, f1001, f997])).
fof(f97, plain, ((e23 = op2(e20, e23)) | (e22 = op2(e20, e23)) | (e21 = op2(e20, e23)) | (e20 = op2(e20, e23))), inference(cnf_transformation, [], [f3])).
fof(f3, plain, (((e23 = op2(e23, e23)) | (e22 = op2(e23, e23)) | (e21 = op2(e23, e23)) | (e20 = op2(e23, e23))) & ((e23 = op2(e23, e22)) | (e22 = op2(e23, e22)) | (e21 = op2(e23, e22)) | (e20 = op2(e23, e22))) & ((e23 = op2(e23, e21)) | (e22 = op2(e23, e21)) | (e21 = op2(e23, e21)) | (e20 = op2(e23, e21))) & ((e23 = op2(e23, e20)) | (e22 = op2(e23, e20)) | (e21 = op2(e23, e20)) | (e20 = op2(e23, e20))) & ((e23 = op2(e22, e23)) | (e22 = op2(e22, e23)) | (e21 = op2(e22, e23)) | (e20 = op2(e22, e23))) & ((e23 = op2(e22, e22)) | (e22 = op2(e22, e22)) | (e21 = op2(e22, e22)) | (e20 = op2(e22, e22))) & ((e23 = op2(e22, e21)) | (e22 = op2(e22, e21)) | (e21 = op2(e22, e21)) | (e20 = op2(e22, e21))) & ((e23 = op2(e22, e20)) | (e22 = op2(e22, e20)) | (e21 = op2(e22, e20)) | (e20 = op2(e22, e20))) & ((e23 = op2(e21, e23)) | (e22 = op2(e21, e23)) | (e21 = op2(e21, e23)) | (e20 = op2(e21, e23))) & ((e23 = op2(e21, e22)) | (e22 = op2(e21, e22)) | (e21 = op2(e21, e22)) | (e20 = op2(e21, e22))) & ((e23 = op2(e21, e21)) | (e22 = op2(e21, e21)) | (e21 = op2(e21, e21)) | (e20 = op2(e21, e21))) & ((e23 = op2(e21, e20)) | (e22 = op2(e21, e20)) | (e21 = op2(e21, e20)) | (e20 = op2(e21, e20))) & ((e23 = op2(e20, e23)) | (e22 = op2(e20, e23)) | (e21 = op2(e20, e23)) | (e20 = op2(e20, e23))) & ((e23 = op2(e20, e22)) | (e22 = op2(e20, e22)) | (e21 = op2(e20, e22)) | (e20 = op2(e20, e22))) & ((e23 = op2(e20, e21)) | (e22 = op2(e20, e21)) | (e21 = op2(e20, e21)) | (e20 = op2(e20, e21))) & ((op2(e20, e20) = e23) | (op2(e20, e20) = e22) | (op2(e20, e20) = e21) | (e20 = op2(e20, e20)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG111+1.p', ax3)).
fof(f791, plain, (spl12_61 | spl12_57 | spl12_53 | spl12_49), inference(avatar_split_clause, [], [f62, f693, f710, f727, f744])).
fof(f62, plain, ((e10 = op1(e10, e13)) | (e10 = op1(e10, e12)) | (e10 = op1(e10, e11)) | (e10 = op1(e10, e10))), inference(cnf_transformation, [], [f2])).
fof(f2, plain, (((e13 = op1(e13, e13)) | (e13 = op1(e12, e13)) | (e13 = op1(e11, e13)) | (e13 = op1(e10, e13))) & ((e13 = op1(e13, e13)) | (e13 = op1(e13, e12)) | (e13 = op1(e13, e11)) | (e13 = op1(e13, e10))) & ((e12 = op1(e13, e13)) | (e12 = op1(e12, e13)) | (e12 = op1(e11, e13)) | (e12 = op1(e10, e13))) & ((e12 = op1(e13, e13)) | (e12 = op1(e13, e12)) | (e12 = op1(e13, e11)) | (e12 = op1(e13, e10))) & ((e11 = op1(e13, e13)) | (e11 = op1(e12, e13)) | (e11 = op1(e11, e13)) | (e11 = op1(e10, e13))) & ((e11 = op1(e13, e13)) | (e11 = op1(e13, e12)) | (e11 = op1(e13, e11)) | (e11 = op1(e13, e10))) & ((e10 = op1(e13, e13)) | (e10 = op1(e12, e13)) | (e10 = op1(e11, e13)) | (e10 = op1(e10, e13))) & ((e10 = op1(e13, e13)) | (e10 = op1(e13, e12)) | (e10 = op1(e13, e11)) | (e10 = op1(e13, e10))) & ((e13 = op1(e13, e12)) | (e13 = op1(e12, e12)) | (e13 = op1(e11, e12)) | (e13 = op1(e10, e12))) & ((e13 = op1(e12, e13)) | (e13 = op1(e12, e12)) | (e13 = op1(e12, e11)) | (e13 = op1(e12, e10))) & ((e12 = op1(e13, e12)) | (e12 = op1(e12, e12)) | (e12 = op1(e11, e12)) | (e12 = op1(e10, e12))) & ((e12 = op1(e12, e13)) | (e12 = op1(e12, e12)) | (e12 = op1(e12, e11)) | (e12 = op1(e12, e10))) & ((e11 = op1(e13, e12)) | (e11 = op1(e12, e12)) | (e11 = op1(e11, e12)) | (e11 = op1(e10, e12))) & ((e11 = op1(e12, e13)) | (e11 = op1(e12, e12)) | (e11 = op1(e12, e11)) | (e11 = op1(e12, e10))) & ((e10 = op1(e13, e12)) | (e10 = op1(e12, e12)) | (e10 = op1(e11, e12)) | (e10 = op1(e10, e12))) & ((e10 = op1(e12, e13)) | (e10 = op1(e12, e12)) | (e10 = op1(e12, e11)) | (e10 = op1(e12, e10))) & ((e13 = op1(e13, e11)) | (e13 = op1(e12, e11)) | (e13 = op1(e11, e11)) | (e13 = op1(e10, e11))) & ((e13 = op1(e11, e13)) | (e13 = op1(e11, e12)) | (e13 = op1(e11, e11)) | (e13 = op1(e11, e10))) & ((e12 = op1(e13, e11)) | (e12 = op1(e12, e11)) | (e12 = op1(e11, e11)) | (e12 = op1(e10, e11))) & ((e12 = op1(e11, e13)) | (e12 = op1(e11, e12)) | (e12 = op1(e11, e11)) | (e12 = op1(e11, e10))) & ((e11 = op1(e13, e11)) | (e11 = op1(e12, e11)) | (e11 = op1(e11, e11)) | (e11 = op1(e10, e11))) & ((e11 = op1(e11, e13)) | (e11 = op1(e11, e12)) | (e11 = op1(e11, e11)) | (e11 = op1(e11, e10))) & ((e10 = op1(e13, e11)) | (e10 = op1(e12, e11)) | (e10 = op1(e11, e11)) | (e10 = op1(e10, e11))) & ((e10 = op1(e11, e13)) | (e10 = op1(e11, e12)) | (e10 = op1(e11, e11)) | (e10 = op1(e11, e10))) & ((e13 = op1(e13, e10)) | (e13 = op1(e12, e10)) | (e13 = op1(e11, e10)) | (op1(e10, e10) = e13)) & ((e13 = op1(e10, e13)) | (e13 = op1(e10, e12)) | (e13 = op1(e10, e11)) | (op1(e10, e10) = e13)) & ((e12 = op1(e13, e10)) | (e12 = op1(e12, e10)) | (e12 = op1(e11, e10)) | (op1(e10, e10) = e12)) & ((e12 = op1(e10, e13)) | (e12 = op1(e10, e12)) | (e12 = op1(e10, e11)) | (op1(e10, e10) = e12)) & ((e11 = op1(e13, e10)) | (e11 = op1(e12, e10)) | (e11 = op1(e11, e10)) | (op1(e10, e10) = e11)) & ((e11 = op1(e10, e13)) | (e11 = op1(e10, e12)) | (e11 = op1(e10, e11)) | (op1(e10, e10) = e11)) & ((e10 = op1(e13, e10)) | (e10 = op1(e12, e10)) | (e10 = op1(e11, e10)) | (e10 = op1(e10, e10))) & ((e10 = op1(e10, e13)) | (e10 = op1(e10, e12)) | (e10 = op1(e10, e11)) | (e10 = op1(e10, e10)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG111+1.p', ax2)).
fof(f783, plain, (spl12_45 | spl12_41 | spl12_37 | spl12_33), inference(avatar_split_clause, [], [f70, f625, f642, f659, f676])).
fof(f70, plain, ((e10 = op1(e11, e13)) | (e10 = op1(e11, e12)) | (e10 = op1(e11, e11)) | (e10 = op1(e11, e10))), inference(cnf_transformation, [], [f2])).
fof(f781, plain, (spl12_46 | spl12_42 | spl12_38 | spl12_34), inference(avatar_split_clause, [], [f72, f629, f646, f663, f680])).
fof(f72, plain, ((e11 = op1(e11, e13)) | (e11 = op1(e11, e12)) | (e11 = op1(e11, e11)) | (e11 = op1(e11, e10))), inference(cnf_transformation, [], [f2])).
fof(f780, plain, (spl12_58 | spl12_42 | spl12_26 | spl12_10), inference(avatar_split_clause, [], [f73, f527, f595, f663, f731])).
fof(f73, plain, ((e11 = op1(e13, e11)) | (e11 = op1(e12, e11)) | (e11 = op1(e11, e11)) | (e11 = op1(e10, e11))), inference(cnf_transformation, [], [f2])).
fof(f775, plain, (spl12_29 | spl12_25 | spl12_21 | spl12_17), inference(avatar_split_clause, [], [f78, f557, f574, f591, f608])).
fof(f78, plain, ((e10 = op1(e12, e13)) | (e10 = op1(e12, e12)) | (e10 = op1(e12, e11)) | (e10 = op1(e12, e10))), inference(cnf_transformation, [], [f2])).
fof(f771, plain, (spl12_31 | spl12_27 | spl12_23 | spl12_19), inference(avatar_split_clause, [], [f82, f565, f582, f599, f616])).
fof(f82, plain, ((e12 = op1(e12, e13)) | (e12 = op1(e12, e12)) | (e12 = op1(e12, e11)) | (e12 = op1(e12, e10))), inference(cnf_transformation, [], [f2])).
fof(f768, plain, (spl12_56 | spl12_40 | spl12_24 | spl12_8), inference(avatar_split_clause, [], [f85, f518, f586, f654, f722])).
fof(f85, plain, ((e13 = op1(e13, e12)) | (e13 = op1(e12, e12)) | (e13 = op1(e11, e12)) | (e13 = op1(e10, e12))), inference(cnf_transformation, [], [f2])).
fof(f763, plain, (spl12_15 | spl12_11 | spl12_7 | spl12_3), inference(avatar_split_clause, [], [f90, f497, f514, f531, f548])).
fof(f90, plain, ((e12 = op1(e13, e13)) | (e12 = op1(e13, e12)) | (e12 = op1(e13, e11)) | (e12 = op1(e13, e10))), inference(cnf_transformation, [], [f2])).
fof(f762, plain, (spl12_51 | spl12_35 | spl12_19 | spl12_3), inference(avatar_split_clause, [], [f91, f497, f565, f633, f701])).
fof(f91, plain, ((e12 = op1(e13, e13)) | (e12 = op1(e12, e13)) | (e12 = op1(e11, e13)) | (e12 = op1(e10, e13))), inference(cnf_transformation, [], [f2])).
fof(f761, plain, (spl12_16 | spl12_12 | spl12_8 | spl12_4), inference(avatar_split_clause, [], [f92, f501, f518, f535, f552])).
fof(f92, plain, ((e13 = op1(e13, e13)) | (e13 = op1(e13, e12)) | (e13 = op1(e13, e11)) | (e13 = op1(e13, e10))), inference(cnf_transformation, [], [f2])).
fof(f759, plain, (spl12_61 | spl12_62 | spl12_63 | spl12_64), inference(avatar_split_clause, [], [f46, f756, f752, f748, f744])).
fof(f46, plain, ((op1(e10, e10) = e13) | (op1(e10, e10) = e12) | (op1(e10, e10) = e11) | (e10 = op1(e10, e10))), inference(cnf_transformation, [], [f1])).
fof(f1, plain, (((e13 = op1(e13, e13)) | (e12 = op1(e13, e13)) | (e11 = op1(e13, e13)) | (e10 = op1(e13, e13))) & ((e13 = op1(e13, e12)) | (e12 = op1(e13, e12)) | (e11 = op1(e13, e12)) | (e10 = op1(e13, e12))) & ((e13 = op1(e13, e11)) | (e12 = op1(e13, e11)) | (e11 = op1(e13, e11)) | (e10 = op1(e13, e11))) & ((e13 = op1(e13, e10)) | (e12 = op1(e13, e10)) | (e11 = op1(e13, e10)) | (e10 = op1(e13, e10))) & ((e13 = op1(e12, e13)) | (e12 = op1(e12, e13)) | (e11 = op1(e12, e13)) | (e10 = op1(e12, e13))) & ((e13 = op1(e12, e12)) | (e12 = op1(e12, e12)) | (e11 = op1(e12, e12)) | (e10 = op1(e12, e12))) & ((e13 = op1(e12, e11)) | (e12 = op1(e12, e11)) | (e11 = op1(e12, e11)) | (e10 = op1(e12, e11))) & ((e13 = op1(e12, e10)) | (e12 = op1(e12, e10)) | (e11 = op1(e12, e10)) | (e10 = op1(e12, e10))) & ((e13 = op1(e11, e13)) | (e12 = op1(e11, e13)) | (e11 = op1(e11, e13)) | (e10 = op1(e11, e13))) & ((e13 = op1(e11, e12)) | (e12 = op1(e11, e12)) | (e11 = op1(e11, e12)) | (e10 = op1(e11, e12))) & ((e13 = op1(e11, e11)) | (e12 = op1(e11, e11)) | (e11 = op1(e11, e11)) | (e10 = op1(e11, e11))) & ((e13 = op1(e11, e10)) | (e12 = op1(e11, e10)) | (e11 = op1(e11, e10)) | (e10 = op1(e11, e10))) & ((e13 = op1(e10, e13)) | (e12 = op1(e10, e13)) | (e11 = op1(e10, e13)) | (e10 = op1(e10, e13))) & ((e13 = op1(e10, e12)) | (e12 = op1(e10, e12)) | (e11 = op1(e10, e12)) | (e10 = op1(e10, e12))) & ((e13 = op1(e10, e11)) | (e12 = op1(e10, e11)) | (e11 = op1(e10, e11)) | (e10 = op1(e10, e11))) & ((op1(e10, e10) = e13) | (op1(e10, e10) = e12) | (op1(e10, e10) = e11) | (e10 = op1(e10, e10)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG111+1.p', ax1)).
fof(f708, plain, (spl12_49 | spl12_50 | spl12_51 | spl12_52), inference(avatar_split_clause, [], [f49, f705, f701, f697, f693])).
fof(f49, plain, ((e13 = op1(e10, e13)) | (e12 = op1(e10, e13)) | (e11 = op1(e10, e13)) | (e10 = op1(e10, e13))), inference(cnf_transformation, [], [f1])).
fof(f674, plain, (spl12_41 | spl12_42 | spl12_43 | spl12_44), inference(avatar_split_clause, [], [f51, f671, f667, f663, f659])).
fof(f51, plain, ((e13 = op1(e11, e11)) | (e12 = op1(e11, e11)) | (e11 = op1(e11, e11)) | (e10 = op1(e11, e11))), inference(cnf_transformation, [], [f1])).