fof(f983, plain, $false, inference(subsumption_resolution, [], [f982, f165])).
fof(f165, plain, (e0 = op(e0, e0)), inference(cnf_transformation, [], [f2])).
fof(f2, plain, ((e4 = op(e4, e4)) & (e0 = op(e4, e3)) & (e1 = op(e4, e2)) & (e2 = op(e4, e1)) & (e3 = op(e4, e0)) & (e1 = op(e3, e4)) & (e3 = op(e3, e3)) & (e0 = op(e3, e2)) & (e4 = op(e3, e1)) & (e2 = op(e3, e0)) & (e3 = op(e2, e4)) & (e4 = op(e2, e3)) & (e2 = op(e2, e2)) & (e0 = op(e2, e1)) & (e1 = op(e2, e0)) & (e0 = op(e1, e4)) & (e2 = op(e1, e3)) & (e3 = op(e1, e2)) & (e1 = op(e1, e1)) & (e4 = op(e1, e0)) & (e2 = op(e0, e4)) & (e1 = op(e0, e3)) & (e4 = op(e0, e2)) & (e3 = op(e0, e1)) & (e0 = op(e0, e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG188+1.p', ax2)).
fof(f982, plain, ~ (e0 = op(e0, e0)), inference(forward_demodulation, [], [f981, f165])).
fof(f981, plain, ~ (e0 = op(e0, op(e0, e0))), inference(subsumption_resolution, [], [f980, f168])).
fof(f168, plain, (e1 = op(e0, e3)), inference(cnf_transformation, [], [f2])).
fof(f980, plain, (~ (e1 = op(e0, e3)) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f979, f166])).
fof(f166, plain, (e3 = op(e0, e1)), inference(cnf_transformation, [], [f2])).
fof(f979, plain, (~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f978, f169])).
fof(f169, plain, (e2 = op(e0, e4)), inference(cnf_transformation, [], [f2])).
fof(f978, plain, (~ (e2 = op(e0, e4)) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f977, f167])).
fof(f167, plain, (e4 = op(e0, e2)), inference(cnf_transformation, [], [f2])).
fof(f977, plain, (~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f976, f166])).
fof(f976, plain, (~ (e3 = op(e0, e1)) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f975, f168])).
fof(f975, plain, (~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f974, f167])).
fof(f974, plain, (~ (e4 = op(e0, e2)) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f973, f169])).
fof(f973, plain, (~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f972, f174])).
fof(f174, plain, (e0 = op(e1, e4)), inference(cnf_transformation, [], [f2])).
fof(f972, plain, (~ (e0 = op(e1, e4)) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f971, f170])).
fof(f170, plain, (e4 = op(e1, e0)), inference(cnf_transformation, [], [f2])).
fof(f971, plain, (~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f970, f171])).
fof(f171, plain, (e1 = op(e1, e1)), inference(cnf_transformation, [], [f2])).
fof(f970, plain, (~ (e1 = op(e1, e1)) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f969, f171])).
fof(f969, plain, (~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f968, f173])).
fof(f173, plain, (e2 = op(e1, e3)), inference(cnf_transformation, [], [f2])).
fof(f968, plain, (~ (e2 = op(e1, e3)) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f967, f172])).
fof(f172, plain, (e3 = op(e1, e2)), inference(cnf_transformation, [], [f2])).
fof(f967, plain, (~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f966, f172])).
fof(f966, plain, (~ (e3 = op(e1, e2)) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f965, f173])).
fof(f965, plain, (~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f964, f170])).
fof(f964, plain, (~ (e4 = op(e1, e0)) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f963, f174])).
fof(f963, plain, (~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f962, f176])).
fof(f176, plain, (e0 = op(e2, e1)), inference(cnf_transformation, [], [f2])).
fof(f962, plain, (~ (e0 = op(e2, e1)) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f961, f175])).
fof(f175, plain, (e1 = op(e2, e0)), inference(cnf_transformation, [], [f2])).
fof(f961, plain, (~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f960, f175])).
fof(f960, plain, (~ (e1 = op(e2, e0)) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f959, f176])).
fof(f959, plain, (~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f958, f177])).
fof(f177, plain, (e2 = op(e2, e2)), inference(cnf_transformation, [], [f2])).
fof(f958, plain, (~ (e2 = op(e2, e2)) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f957, f177])).
fof(f957, plain, (~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f956, f179])).
fof(f179, plain, (e3 = op(e2, e4)), inference(cnf_transformation, [], [f2])).
fof(f956, plain, (~ (e3 = op(e2, e4)) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f955, f178])).
fof(f178, plain, (e4 = op(e2, e3)), inference(cnf_transformation, [], [f2])).
fof(f955, plain, (~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f954, f178])).
fof(f954, plain, (~ (e4 = op(e2, e3)) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f953, f179])).
fof(f953, plain, (~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f952, f182])).
fof(f182, plain, (e0 = op(e3, e2)), inference(cnf_transformation, [], [f2])).
fof(f952, plain, (~ (e0 = op(e3, e2)) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f951, f180])).
fof(f180, plain, (e2 = op(e3, e0)), inference(cnf_transformation, [], [f2])).
fof(f951, plain, (~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f950, f184])).
fof(f184, plain, (e1 = op(e3, e4)), inference(cnf_transformation, [], [f2])).
fof(f950, plain, (~ (e1 = op(e3, e4)) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f949, f181])).
fof(f181, plain, (e4 = op(e3, e1)), inference(cnf_transformation, [], [f2])).
fof(f949, plain, (~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f948, f180])).
fof(f948, plain, (~ (e2 = op(e3, e0)) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f947, f182])).
fof(f947, plain, (~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f946, f183])).
fof(f183, plain, (e3 = op(e3, e3)), inference(cnf_transformation, [], [f2])).
fof(f946, plain, (~ (e3 = op(e3, e3)) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f945, f183])).
fof(f945, plain, (~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f944, f181])).
fof(f944, plain, (~ (e4 = op(e3, e1)) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f943, f184])).
fof(f943, plain, (~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f942, f188])).
fof(f188, plain, (e0 = op(e4, e3)), inference(cnf_transformation, [], [f2])).
fof(f942, plain, (~ (e0 = op(e4, e3)) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f941, f185])).
fof(f185, plain, (e3 = op(e4, e0)), inference(cnf_transformation, [], [f2])).
fof(f941, plain, (~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f940, f187])).
fof(f187, plain, (e1 = op(e4, e2)), inference(cnf_transformation, [], [f2])).
fof(f940, plain, (~ (e1 = op(e4, e2)) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f939, f186])).
fof(f186, plain, (e2 = op(e4, e1)), inference(cnf_transformation, [], [f2])).
fof(f939, plain, (~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f938, f186])).
fof(f938, plain, (~ (e2 = op(e4, e1)) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f937, f187])).
fof(f937, plain, (~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f936, f185])).
fof(f936, plain, (~ (e3 = op(e4, e0)) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f935, f188])).
fof(f935, plain, (~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f934, f189])).
fof(f189, plain, (e4 = op(e4, e4)), inference(cnf_transformation, [], [f2])).
fof(f934, plain, (~ (e4 = op(e4, e4)) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f933, f189])).
fof(f933, plain, (~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f932, f165])).
fof(f932, plain, (~ (e0 = op(e0, e0)) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f931, f165])).
fof(f931, plain, (~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f930, f188])).
fof(f930, plain, (~ (e0 = op(e4, e3)) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f929, f170])).
fof(f929, plain, (~ (e0 = op(op(e1, e0), e3)) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f928, f166])).
fof(f928, plain, (~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f927, f174])).
fof(f927, plain, (~ (e0 = op(e1, e4)) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f926, f175])).
fof(f926, plain, (~ (e0 = op(op(e2, e0), e4)) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f925, f167])).
fof(f925, plain, (~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f924, f176])).
fof(f924, plain, (~ (e0 = op(e2, e1)) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f923, f180])).
fof(f923, plain, (~ (e0 = op(op(e3, e0), e1)) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f922, f168])).
fof(f922, plain, (~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f921, f182])).
fof(f921, plain, (~ (e0 = op(e3, e2)) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f920, f185])).
fof(f920, plain, (~ (e0 = op(op(e4, e0), e2)) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f919, f169])).
fof(f919, plain, (~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f918, f184])).
fof(f918, plain, (~ (e1 = op(e3, e4)) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f917, f166])).
fof(f917, plain, (~ (e1 = op(op(e0, e1), e4)) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f916, f170])).
fof(f916, plain, (~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f915, f171])).
fof(f915, plain, (~ (e1 = op(e1, e1)) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f914, f171])).
fof(f914, plain, (~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f913, f168])).
fof(f913, plain, (~ (e1 = op(e0, e3)) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f912, f176])).
fof(f912, plain, (~ (e1 = op(op(e2, e1), e3)) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f911, f172])).
fof(f911, plain, (~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f910, f187])).
fof(f910, plain, (~ (e1 = op(e4, e2)) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f909, f181])).
fof(f909, plain, (~ (e1 = op(op(e3, e1), e2)) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f908, f173])).
fof(f908, plain, (~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f907, f175])).
fof(f907, plain, (~ (e1 = op(e2, e0)) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f906, f186])).
fof(f906, plain, (~ (e1 = op(op(e4, e1), e0)) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f905, f174])).
fof(f905, plain, (~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f904, f186])).
fof(f904, plain, (~ (e2 = op(e4, e1)) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f903, f167])).
fof(f903, plain, (~ (e2 = op(op(e0, e2), e1)) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f902, f175])).
fof(f902, plain, (~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f901, f180])).
fof(f901, plain, (~ (e2 = op(e3, e0)) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f900, f172])).
fof(f900, plain, (~ (e2 = op(op(e1, e2), e0)) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f899, f176])).
fof(f899, plain, (~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f898, f177])).
fof(f898, plain, (~ (e2 = op(e2, e2)) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f897, f177])).
fof(f897, plain, (~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f896, f169])).
fof(f896, plain, (~ (e2 = op(e0, e4)) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f895, f182])).
fof(f895, plain, (~ (e2 = op(op(e3, e2), e4)) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f894, f178])).
fof(f894, plain, (~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f893, f173])).
fof(f893, plain, (~ (e2 = op(e1, e3)) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f892, f187])).
fof(f892, plain, (~ (e2 = op(op(e4, e2), e3)) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f891, f179])).
fof(f891, plain, (~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f890, f172])).
fof(f890, plain, (~ (e3 = op(e1, e2)) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f889, f168])).
fof(f889, plain, (~ (e3 = op(op(e0, e3), e2)) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f888, f180])).
fof(f888, plain, (~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f887, f179])).
fof(f887, plain, (~ (e3 = op(e2, e4)) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f886, f173])).
fof(f886, plain, (~ (e3 = op(op(e1, e3), e4)) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f885, f181])).
fof(f885, plain, (~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f884, f185])).
fof(f884, plain, (~ (e3 = op(e4, e0)) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f883, f178])).
fof(f883, plain, (~ (e3 = op(op(e2, e3), e0)) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f882, f182])).
fof(f882, plain, (~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f881, f183])).
fof(f881, plain, (~ (e3 = op(e3, e3)) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f880, f183])).
fof(f880, plain, (~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f879, f166])).
fof(f879, plain, (~ (e3 = op(e0, e1)) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f878, f188])).
fof(f878, plain, (~ (e3 = op(op(e4, e3), e1)) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f877, f184])).
fof(f877, plain, (~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f876, f178])).
fof(f876, plain, (~ (e4 = op(e2, e3)) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f875, f169])).
fof(f875, plain, (~ (e4 = op(op(e0, e4), e3)) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f874, f185])).
fof(f874, plain, (~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f873, f167])).
fof(f873, plain, (~ (e4 = op(e0, e2)) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f872, f174])).
fof(f872, plain, (~ (e4 = op(op(e1, e4), e2)) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f871, f186])).
fof(f871, plain, (~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f870, f181])).
fof(f870, plain, (~ (e4 = op(e3, e1)) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f869, f179])).
fof(f869, plain, (~ (e4 = op(op(e2, e4), e1)) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f868, f187])).
fof(f868, plain, (~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f867, f170])).
fof(f867, plain, (~ (e4 = op(e1, e0)) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f866, f184])).
fof(f866, plain, (~ (e4 = op(op(e3, e4), e0)) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f865, f188])).
fof(f865, plain, (~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f864, f189])).
fof(f864, plain, (~ (e4 = op(e4, e4)) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(forward_demodulation, [], [f863, f189])).
fof(f863, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f862, f788])).
fof(f788, plain, ~ sP0, inference(subsumption_resolution, [], [f555, f165])).
fof(f555, plain, (~ (e0 = op(e0, e0)) | ~ sP0), inference(cnf_transformation, [], [f154])).
fof(f154, plain, ((~ (e4 = op(e0, e0)) & ~ (e3 = op(e0, e0)) & ~ (e2 = op(e0, e0)) & ~ (e1 = op(e0, e0)) & ~ (e0 = op(e0, e0))) | ~ sP0), inference(nnf_transformation, [], [f6])).
fof(f6, plain, ((~ (e4 = op(e0, e0)) & ~ (e3 = op(e0, e0)) & ~ (e2 = op(e0, e0)) & ~ (e1 = op(e0, e0)) & ~ (e0 = op(e0, e0))) | ~ sP0), inference(usedef, [], [e6])).
fof(e6, plain, (sP0 <=> (~ (e4 = op(e0, e0)) & ~ (e3 = op(e0, e0)) & ~ (e2 = op(e0, e0)) & ~ (e1 = op(e0, e0)) & ~ (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f862, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f861, f783])).
fof(f783, plain, ~ sP1, inference(subsumption_resolution, [], [f553, f166])).
fof(f553, plain, (~ (e3 = op(e0, e1)) | ~ sP1), inference(cnf_transformation, [], [f153])).
fof(f153, plain, ((~ (e4 = op(e0, e1)) & ~ (e3 = op(e0, e1)) & ~ (e2 = op(e0, e1)) & ~ (e1 = op(e0, e1)) & ~ (e0 = op(e0, e1))) | ~ sP1), inference(nnf_transformation, [], [f7])).
fof(f7, plain, ((~ (e4 = op(e0, e1)) & ~ (e3 = op(e0, e1)) & ~ (e2 = op(e0, e1)) & ~ (e1 = op(e0, e1)) & ~ (e0 = op(e0, e1))) | ~ sP1), inference(usedef, [], [e7])).
fof(e7, plain, (sP1 <=> (~ (e4 = op(e0, e1)) & ~ (e3 = op(e0, e1)) & ~ (e2 = op(e0, e1)) & ~ (e1 = op(e0, e1)) & ~ (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f861, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f860, f781])).
fof(f781, plain, ~ sP2, inference(subsumption_resolution, [], [f549, f167])).
fof(f549, plain, (~ (e4 = op(e0, e2)) | ~ sP2), inference(cnf_transformation, [], [f152])).
fof(f152, plain, ((~ (e4 = op(e0, e2)) & ~ (e3 = op(e0, e2)) & ~ (e2 = op(e0, e2)) & ~ (e1 = op(e0, e2)) & ~ (e0 = op(e0, e2))) | ~ sP2), inference(nnf_transformation, [], [f8])).
fof(f8, plain, ((~ (e4 = op(e0, e2)) & ~ (e3 = op(e0, e2)) & ~ (e2 = op(e0, e2)) & ~ (e1 = op(e0, e2)) & ~ (e0 = op(e0, e2))) | ~ sP2), inference(usedef, [], [e8])).
fof(e8, plain, (sP2 <=> (~ (e4 = op(e0, e2)) & ~ (e3 = op(e0, e2)) & ~ (e2 = op(e0, e2)) & ~ (e1 = op(e0, e2)) & ~ (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f860, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f859, f780])).
fof(f780, plain, ~ sP3, inference(subsumption_resolution, [], [f541, f168])).
fof(f541, plain, (~ (e1 = op(e0, e3)) | ~ sP3), inference(cnf_transformation, [], [f151])).
fof(f151, plain, ((~ (e4 = op(e0, e3)) & ~ (e3 = op(e0, e3)) & ~ (e2 = op(e0, e3)) & ~ (e1 = op(e0, e3)) & ~ (e0 = op(e0, e3))) | ~ sP3), inference(nnf_transformation, [], [f9])).
fof(f9, plain, ((~ (e4 = op(e0, e3)) & ~ (e3 = op(e0, e3)) & ~ (e2 = op(e0, e3)) & ~ (e1 = op(e0, e3)) & ~ (e0 = op(e0, e3))) | ~ sP3), inference(usedef, [], [e9])).
fof(e9, plain, (sP3 <=> (~ (e4 = op(e0, e3)) & ~ (e3 = op(e0, e3)) & ~ (e2 = op(e0, e3)) & ~ (e1 = op(e0, e3)) & ~ (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f859, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f858, f776])).
fof(f776, plain, ~ sP4, inference(subsumption_resolution, [], [f537, f169])).
fof(f537, plain, (~ (e2 = op(e0, e4)) | ~ sP4), inference(cnf_transformation, [], [f150])).
fof(f150, plain, ((~ (e4 = op(e0, e4)) & ~ (e3 = op(e0, e4)) & ~ (e2 = op(e0, e4)) & ~ (e1 = op(e0, e4)) & ~ (e0 = op(e0, e4))) | ~ sP4), inference(nnf_transformation, [], [f10])).
fof(f10, plain, ((~ (e4 = op(e0, e4)) & ~ (e3 = op(e0, e4)) & ~ (e2 = op(e0, e4)) & ~ (e1 = op(e0, e4)) & ~ (e0 = op(e0, e4))) | ~ sP4), inference(usedef, [], [e10])).
fof(e10, plain, (sP4 <=> (~ (e4 = op(e0, e4)) & ~ (e3 = op(e0, e4)) & ~ (e2 = op(e0, e4)) & ~ (e1 = op(e0, e4)) & ~ (e0 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f858, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f857, f773])).
fof(f773, plain, ~ sP5, inference(subsumption_resolution, [], [f534, f170])).
fof(f534, plain, (~ (e4 = op(e1, e0)) | ~ sP5), inference(cnf_transformation, [], [f149])).
fof(f149, plain, ((~ (e4 = op(e1, e0)) & ~ (e3 = op(e1, e0)) & ~ (e2 = op(e1, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e1, e0))) | ~ sP5), inference(nnf_transformation, [], [f11])).
fof(f11, plain, ((~ (e4 = op(e1, e0)) & ~ (e3 = op(e1, e0)) & ~ (e2 = op(e1, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e1, e0))) | ~ sP5), inference(usedef, [], [e11])).
fof(e11, plain, (sP5 <=> (~ (e4 = op(e1, e0)) & ~ (e3 = op(e1, e0)) & ~ (e2 = op(e1, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f857, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f856, f772])).
fof(f772, plain, ~ sP6, inference(subsumption_resolution, [], [f526, f171])).
fof(f526, plain, (~ (e1 = op(e1, e1)) | ~ sP6), inference(cnf_transformation, [], [f148])).
fof(f148, plain, ((~ (e4 = op(e1, e1)) & ~ (e3 = op(e1, e1)) & ~ (e2 = op(e1, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e1, e1))) | ~ sP6), inference(nnf_transformation, [], [f12])).
fof(f12, plain, ((~ (e4 = op(e1, e1)) & ~ (e3 = op(e1, e1)) & ~ (e2 = op(e1, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e1, e1))) | ~ sP6), inference(usedef, [], [e12])).
fof(e12, plain, (sP6 <=> (~ (e4 = op(e1, e1)) & ~ (e3 = op(e1, e1)) & ~ (e2 = op(e1, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f856, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f855, f768])).
fof(f768, plain, ~ sP7, inference(subsumption_resolution, [], [f523, f172])).
fof(f523, plain, (~ (e3 = op(e1, e2)) | ~ sP7), inference(cnf_transformation, [], [f147])).
fof(f147, plain, ((~ (e4 = op(e1, e2)) & ~ (e3 = op(e1, e2)) & ~ (e2 = op(e1, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e1, e2))) | ~ sP7), inference(nnf_transformation, [], [f13])).
fof(f13, plain, ((~ (e4 = op(e1, e2)) & ~ (e3 = op(e1, e2)) & ~ (e2 = op(e1, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e1, e2))) | ~ sP7), inference(usedef, [], [e13])).
fof(e13, plain, (sP7 <=> (~ (e4 = op(e1, e2)) & ~ (e3 = op(e1, e2)) & ~ (e2 = op(e1, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f855, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f854, f766])).
fof(f766, plain, ~ sP8, inference(subsumption_resolution, [], [f517, f173])).
fof(f517, plain, (~ (e2 = op(e1, e3)) | ~ sP8), inference(cnf_transformation, [], [f146])).
fof(f146, plain, ((~ (e4 = op(e1, e3)) & ~ (e3 = op(e1, e3)) & ~ (e2 = op(e1, e3)) & ~ (e1 = op(e1, e3)) & ~ (e0 = op(e1, e3))) | ~ sP8), inference(nnf_transformation, [], [f14])).
fof(f14, plain, ((~ (e4 = op(e1, e3)) & ~ (e3 = op(e1, e3)) & ~ (e2 = op(e1, e3)) & ~ (e1 = op(e1, e3)) & ~ (e0 = op(e1, e3))) | ~ sP8), inference(usedef, [], [e14])).
fof(e14, plain, (sP8 <=> (~ (e4 = op(e1, e3)) & ~ (e3 = op(e1, e3)) & ~ (e2 = op(e1, e3)) & ~ (e1 = op(e1, e3)) & ~ (e0 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f854, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f853, f763])).
fof(f763, plain, ~ sP9, inference(subsumption_resolution, [], [f510, f174])).
fof(f510, plain, (~ (e0 = op(e1, e4)) | ~ sP9), inference(cnf_transformation, [], [f145])).
fof(f145, plain, ((~ (e4 = op(e1, e4)) & ~ (e3 = op(e1, e4)) & ~ (e2 = op(e1, e4)) & ~ (e1 = op(e1, e4)) & ~ (e0 = op(e1, e4))) | ~ sP9), inference(nnf_transformation, [], [f15])).
fof(f15, plain, ((~ (e4 = op(e1, e4)) & ~ (e3 = op(e1, e4)) & ~ (e2 = op(e1, e4)) & ~ (e1 = op(e1, e4)) & ~ (e0 = op(e1, e4))) | ~ sP9), inference(usedef, [], [e15])).
fof(e15, plain, (sP9 <=> (~ (e4 = op(e1, e4)) & ~ (e3 = op(e1, e4)) & ~ (e2 = op(e1, e4)) & ~ (e1 = op(e1, e4)) & ~ (e0 = op(e1, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f853, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f852, f758])).
fof(f758, plain, ~ sP10, inference(subsumption_resolution, [], [f506, f175])).
fof(f506, plain, (~ (e1 = op(e2, e0)) | ~ sP10), inference(cnf_transformation, [], [f144])).
fof(f144, plain, ((~ (e4 = op(e2, e0)) & ~ (e3 = op(e2, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e2, e0)) & ~ (e0 = op(e2, e0))) | ~ sP10), inference(nnf_transformation, [], [f16])).
fof(f16, plain, ((~ (e4 = op(e2, e0)) & ~ (e3 = op(e2, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e2, e0)) & ~ (e0 = op(e2, e0))) | ~ sP10), inference(usedef, [], [e16])).
fof(e16, plain, (sP10 <=> (~ (e4 = op(e2, e0)) & ~ (e3 = op(e2, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e2, e0)) & ~ (e0 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f852, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f851, f754])).
fof(f754, plain, ~ sP11, inference(subsumption_resolution, [], [f500, f176])).
fof(f500, plain, (~ (e0 = op(e2, e1)) | ~ sP11), inference(cnf_transformation, [], [f143])).
fof(f143, plain, ((~ (e4 = op(e2, e1)) & ~ (e3 = op(e2, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e2, e1)) & ~ (e0 = op(e2, e1))) | ~ sP11), inference(nnf_transformation, [], [f17])).
fof(f17, plain, ((~ (e4 = op(e2, e1)) & ~ (e3 = op(e2, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e2, e1)) & ~ (e0 = op(e2, e1))) | ~ sP11), inference(usedef, [], [e17])).
fof(e17, plain, (sP11 <=> (~ (e4 = op(e2, e1)) & ~ (e3 = op(e2, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e2, e1)) & ~ (e0 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f851, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f850, f749])).
fof(f749, plain, ~ sP12, inference(subsumption_resolution, [], [f497, f177])).
fof(f497, plain, (~ (e2 = op(e2, e2)) | ~ sP12), inference(cnf_transformation, [], [f142])).
fof(f142, plain, ((~ (e4 = op(e2, e2)) & ~ (e3 = op(e2, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e2, e2)) & ~ (e0 = op(e2, e2))) | ~ sP12), inference(nnf_transformation, [], [f18])).
fof(f18, plain, ((~ (e4 = op(e2, e2)) & ~ (e3 = op(e2, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e2, e2)) & ~ (e0 = op(e2, e2))) | ~ sP12), inference(usedef, [], [e18])).
fof(e18, plain, (sP12 <=> (~ (e4 = op(e2, e2)) & ~ (e3 = op(e2, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e2, e2)) & ~ (e0 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f850, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f849, f746])).
fof(f746, plain, ~ sP13, inference(subsumption_resolution, [], [f494, f178])).
fof(f494, plain, (~ (e4 = op(e2, e3)) | ~ sP13), inference(cnf_transformation, [], [f141])).
fof(f141, plain, ((~ (e4 = op(e2, e3)) & ~ (e3 = op(e2, e3)) & ~ (e2 = op(e2, e3)) & ~ (e1 = op(e2, e3)) & ~ (e0 = op(e2, e3))) | ~ sP13), inference(nnf_transformation, [], [f19])).
fof(f19, plain, ((~ (e4 = op(e2, e3)) & ~ (e3 = op(e2, e3)) & ~ (e2 = op(e2, e3)) & ~ (e1 = op(e2, e3)) & ~ (e0 = op(e2, e3))) | ~ sP13), inference(usedef, [], [e19])).
fof(e19, plain, (sP13 <=> (~ (e4 = op(e2, e3)) & ~ (e3 = op(e2, e3)) & ~ (e2 = op(e2, e3)) & ~ (e1 = op(e2, e3)) & ~ (e0 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f849, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f848, f745])).
fof(f745, plain, ~ sP14, inference(subsumption_resolution, [], [f488, f179])).
fof(f488, plain, (~ (e3 = op(e2, e4)) | ~ sP14), inference(cnf_transformation, [], [f140])).
fof(f140, plain, ((~ (e4 = op(e2, e4)) & ~ (e3 = op(e2, e4)) & ~ (e2 = op(e2, e4)) & ~ (e1 = op(e2, e4)) & ~ (e0 = op(e2, e4))) | ~ sP14), inference(nnf_transformation, [], [f20])).
fof(f20, plain, ((~ (e4 = op(e2, e4)) & ~ (e3 = op(e2, e4)) & ~ (e2 = op(e2, e4)) & ~ (e1 = op(e2, e4)) & ~ (e0 = op(e2, e4))) | ~ sP14), inference(usedef, [], [e20])).
fof(e20, plain, (sP14 <=> (~ (e4 = op(e2, e4)) & ~ (e3 = op(e2, e4)) & ~ (e2 = op(e2, e4)) & ~ (e1 = op(e2, e4)) & ~ (e0 = op(e2, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f848, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f847, f743])).
fof(f743, plain, ~ sP15, inference(subsumption_resolution, [], [f482, f180])).
fof(f482, plain, (~ (e2 = op(e3, e0)) | ~ sP15), inference(cnf_transformation, [], [f139])).
fof(f139, plain, ((~ (e4 = op(e3, e0)) & ~ (e3 = op(e3, e0)) & ~ (e2 = op(e3, e0)) & ~ (e1 = op(e3, e0)) & ~ (e0 = op(e3, e0))) | ~ sP15), inference(nnf_transformation, [], [f21])).
fof(f21, plain, ((~ (e4 = op(e3, e0)) & ~ (e3 = op(e3, e0)) & ~ (e2 = op(e3, e0)) & ~ (e1 = op(e3, e0)) & ~ (e0 = op(e3, e0))) | ~ sP15), inference(usedef, [], [e21])).
fof(e21, plain, (sP15 <=> (~ (e4 = op(e3, e0)) & ~ (e3 = op(e3, e0)) & ~ (e2 = op(e3, e0)) & ~ (e1 = op(e3, e0)) & ~ (e0 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f847, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f846, f740])).
fof(f740, plain, ~ sP16, inference(subsumption_resolution, [], [f479, f181])).
fof(f479, plain, (~ (e4 = op(e3, e1)) | ~ sP16), inference(cnf_transformation, [], [f138])).
fof(f138, plain, ((~ (e4 = op(e3, e1)) & ~ (e3 = op(e3, e1)) & ~ (e2 = op(e3, e1)) & ~ (e1 = op(e3, e1)) & ~ (e0 = op(e3, e1))) | ~ sP16), inference(nnf_transformation, [], [f22])).
fof(f22, plain, ((~ (e4 = op(e3, e1)) & ~ (e3 = op(e3, e1)) & ~ (e2 = op(e3, e1)) & ~ (e1 = op(e3, e1)) & ~ (e0 = op(e3, e1))) | ~ sP16), inference(usedef, [], [e22])).
fof(e22, plain, (sP16 <=> (~ (e4 = op(e3, e1)) & ~ (e3 = op(e3, e1)) & ~ (e2 = op(e3, e1)) & ~ (e1 = op(e3, e1)) & ~ (e0 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f846, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f845, f739])).
fof(f739, plain, ~ sP17, inference(subsumption_resolution, [], [f470, f182])).
fof(f470, plain, (~ (e0 = op(e3, e2)) | ~ sP17), inference(cnf_transformation, [], [f137])).
fof(f137, plain, ((~ (e4 = op(e3, e2)) & ~ (e3 = op(e3, e2)) & ~ (e2 = op(e3, e2)) & ~ (e1 = op(e3, e2)) & ~ (e0 = op(e3, e2))) | ~ sP17), inference(nnf_transformation, [], [f23])).
fof(f23, plain, ((~ (e4 = op(e3, e2)) & ~ (e3 = op(e3, e2)) & ~ (e2 = op(e3, e2)) & ~ (e1 = op(e3, e2)) & ~ (e0 = op(e3, e2))) | ~ sP17), inference(usedef, [], [e23])).
fof(e23, plain, (sP17 <=> (~ (e4 = op(e3, e2)) & ~ (e3 = op(e3, e2)) & ~ (e2 = op(e3, e2)) & ~ (e1 = op(e3, e2)) & ~ (e0 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f845, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f844, f734])).
fof(f734, plain, ~ sP18, inference(subsumption_resolution, [], [f468, f183])).
fof(f468, plain, (~ (e3 = op(e3, e3)) | ~ sP18), inference(cnf_transformation, [], [f136])).
fof(f136, plain, ((~ (e4 = op(e3, e3)) & ~ (e3 = op(e3, e3)) & ~ (e2 = op(e3, e3)) & ~ (e1 = op(e3, e3)) & ~ (e0 = op(e3, e3))) | ~ sP18), inference(nnf_transformation, [], [f24])).
fof(f24, plain, ((~ (e4 = op(e3, e3)) & ~ (e3 = op(e3, e3)) & ~ (e2 = op(e3, e3)) & ~ (e1 = op(e3, e3)) & ~ (e0 = op(e3, e3))) | ~ sP18), inference(usedef, [], [e24])).
fof(e24, plain, (sP18 <=> (~ (e4 = op(e3, e3)) & ~ (e3 = op(e3, e3)) & ~ (e2 = op(e3, e3)) & ~ (e1 = op(e3, e3)) & ~ (e0 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f844, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f843, f732])).
fof(f732, plain, ~ sP19, inference(subsumption_resolution, [], [f461, f184])).
fof(f461, plain, (~ (e1 = op(e3, e4)) | ~ sP19), inference(cnf_transformation, [], [f135])).
fof(f135, plain, ((~ (e4 = op(e3, e4)) & ~ (e3 = op(e3, e4)) & ~ (e2 = op(e3, e4)) & ~ (e1 = op(e3, e4)) & ~ (e0 = op(e3, e4))) | ~ sP19), inference(nnf_transformation, [], [f25])).
fof(f25, plain, ((~ (e4 = op(e3, e4)) & ~ (e3 = op(e3, e4)) & ~ (e2 = op(e3, e4)) & ~ (e1 = op(e3, e4)) & ~ (e0 = op(e3, e4))) | ~ sP19), inference(usedef, [], [e25])).
fof(e25, plain, (sP19 <=> (~ (e4 = op(e3, e4)) & ~ (e3 = op(e3, e4)) & ~ (e2 = op(e3, e4)) & ~ (e1 = op(e3, e4)) & ~ (e0 = op(e3, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP19])])).
fof(f843, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f842, f728])).
fof(f728, plain, ~ sP20, inference(subsumption_resolution, [], [f458, f185])).
fof(f458, plain, (~ (e3 = op(e4, e0)) | ~ sP20), inference(cnf_transformation, [], [f134])).
fof(f134, plain, ((~ (e4 = op(e4, e0)) & ~ (e3 = op(e4, e0)) & ~ (e2 = op(e4, e0)) & ~ (e1 = op(e4, e0)) & ~ (e0 = op(e4, e0))) | ~ sP20), inference(nnf_transformation, [], [f26])).
fof(f26, plain, ((~ (e4 = op(e4, e0)) & ~ (e3 = op(e4, e0)) & ~ (e2 = op(e4, e0)) & ~ (e1 = op(e4, e0)) & ~ (e0 = op(e4, e0))) | ~ sP20), inference(usedef, [], [e26])).
fof(e26, plain, (sP20 <=> (~ (e4 = op(e4, e0)) & ~ (e3 = op(e4, e0)) & ~ (e2 = op(e4, e0)) & ~ (e1 = op(e4, e0)) & ~ (e0 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP20])])).
fof(f842, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f841, f726])).
fof(f726, plain, ~ sP21, inference(subsumption_resolution, [], [f452, f186])).
fof(f452, plain, (~ (e2 = op(e4, e1)) | ~ sP21), inference(cnf_transformation, [], [f133])).
fof(f133, plain, ((~ (e4 = op(e4, e1)) & ~ (e3 = op(e4, e1)) & ~ (e2 = op(e4, e1)) & ~ (e1 = op(e4, e1)) & ~ (e0 = op(e4, e1))) | ~ sP21), inference(nnf_transformation, [], [f27])).
fof(f27, plain, ((~ (e4 = op(e4, e1)) & ~ (e3 = op(e4, e1)) & ~ (e2 = op(e4, e1)) & ~ (e1 = op(e4, e1)) & ~ (e0 = op(e4, e1))) | ~ sP21), inference(usedef, [], [e27])).
fof(e27, plain, (sP21 <=> (~ (e4 = op(e4, e1)) & ~ (e3 = op(e4, e1)) & ~ (e2 = op(e4, e1)) & ~ (e1 = op(e4, e1)) & ~ (e0 = op(e4, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP21])])).
fof(f841, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f840, f723])).
fof(f723, plain, ~ sP22, inference(subsumption_resolution, [], [f446, f187])).
fof(f446, plain, (~ (e1 = op(e4, e2)) | ~ sP22), inference(cnf_transformation, [], [f132])).
fof(f132, plain, ((~ (e4 = op(e4, e2)) & ~ (e3 = op(e4, e2)) & ~ (e2 = op(e4, e2)) & ~ (e1 = op(e4, e2)) & ~ (e0 = op(e4, e2))) | ~ sP22), inference(nnf_transformation, [], [f28])).
fof(f28, plain, ((~ (e4 = op(e4, e2)) & ~ (e3 = op(e4, e2)) & ~ (e2 = op(e4, e2)) & ~ (e1 = op(e4, e2)) & ~ (e0 = op(e4, e2))) | ~ sP22), inference(usedef, [], [e28])).
fof(e28, plain, (sP22 <=> (~ (e4 = op(e4, e2)) & ~ (e3 = op(e4, e2)) & ~ (e2 = op(e4, e2)) & ~ (e1 = op(e4, e2)) & ~ (e0 = op(e4, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP22])])).
fof(f840, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f839, f719])).
fof(f719, plain, ~ sP23, inference(subsumption_resolution, [], [f440, f188])).
fof(f440, plain, (~ (e0 = op(e4, e3)) | ~ sP23), inference(cnf_transformation, [], [f131])).
fof(f131, plain, ((~ (e4 = op(e4, e3)) & ~ (e3 = op(e4, e3)) & ~ (e2 = op(e4, e3)) & ~ (e1 = op(e4, e3)) & ~ (e0 = op(e4, e3))) | ~ sP23), inference(nnf_transformation, [], [f29])).
fof(f29, plain, ((~ (e4 = op(e4, e3)) & ~ (e3 = op(e4, e3)) & ~ (e2 = op(e4, e3)) & ~ (e1 = op(e4, e3)) & ~ (e0 = op(e4, e3))) | ~ sP23), inference(usedef, [], [e29])).
fof(e29, plain, (sP23 <=> (~ (e4 = op(e4, e3)) & ~ (e3 = op(e4, e3)) & ~ (e2 = op(e4, e3)) & ~ (e1 = op(e4, e3)) & ~ (e0 = op(e4, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP23])])).
fof(f839, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f838, f714])).
fof(f714, plain, ~ sP24, inference(subsumption_resolution, [], [f439, f189])).
fof(f439, plain, (~ (e4 = op(e4, e4)) | ~ sP24), inference(cnf_transformation, [], [f130])).
fof(f130, plain, ((~ (e4 = op(e4, e4)) & ~ (e3 = op(e4, e4)) & ~ (e2 = op(e4, e4)) & ~ (e1 = op(e4, e4)) & ~ (e0 = op(e4, e4))) | ~ sP24), inference(nnf_transformation, [], [f30])).
fof(f30, plain, ((~ (e4 = op(e4, e4)) & ~ (e3 = op(e4, e4)) & ~ (e2 = op(e4, e4)) & ~ (e1 = op(e4, e4)) & ~ (e0 = op(e4, e4))) | ~ sP24), inference(usedef, [], [e30])).
fof(e30, plain, (sP24 <=> (~ (e4 = op(e4, e4)) & ~ (e3 = op(e4, e4)) & ~ (e2 = op(e4, e4)) & ~ (e1 = op(e4, e4)) & ~ (e0 = op(e4, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP24])])).
fof(f838, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f837, f713])).
fof(f713, plain, ~ sP25, inference(subsumption_resolution, [], [f430, f165])).
fof(f430, plain, (~ (e0 = op(e0, e0)) | ~ sP25), inference(cnf_transformation, [], [f129])).
fof(f129, plain, ((~ (e0 = op(e0, e4)) & ~ (e0 = op(e0, e3)) & ~ (e0 = op(e0, e2)) & ~ (e0 = op(e0, e1)) & ~ (e0 = op(e0, e0))) | ~ sP25), inference(nnf_transformation, [], [f31])).
fof(f31, plain, ((~ (e0 = op(e0, e4)) & ~ (e0 = op(e0, e3)) & ~ (e0 = op(e0, e2)) & ~ (e0 = op(e0, e1)) & ~ (e0 = op(e0, e0))) | ~ sP25), inference(usedef, [], [e31])).
fof(e31, plain, (sP25 <=> (~ (e0 = op(e0, e4)) & ~ (e0 = op(e0, e3)) & ~ (e0 = op(e0, e2)) & ~ (e0 = op(e0, e1)) & ~ (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP25])])).
fof(f837, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f836, f708])).
fof(f708, plain, ~ sP26, inference(subsumption_resolution, [], [f425, f165])).
fof(f425, plain, (~ (e0 = op(e0, e0)) | ~ sP26), inference(cnf_transformation, [], [f128])).
fof(f128, plain, ((~ (e0 = op(e4, e0)) & ~ (e0 = op(e3, e0)) & ~ (e0 = op(e2, e0)) & ~ (e0 = op(e1, e0)) & ~ (e0 = op(e0, e0))) | ~ sP26), inference(nnf_transformation, [], [f32])).
fof(f32, plain, ((~ (e0 = op(e4, e0)) & ~ (e0 = op(e3, e0)) & ~ (e0 = op(e2, e0)) & ~ (e0 = op(e1, e0)) & ~ (e0 = op(e0, e0))) | ~ sP26), inference(usedef, [], [e32])).
fof(e32, plain, (sP26 <=> (~ (e0 = op(e4, e0)) & ~ (e0 = op(e3, e0)) & ~ (e0 = op(e2, e0)) & ~ (e0 = op(e1, e0)) & ~ (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP26])])).
fof(f836, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f835, f703])).
fof(f703, plain, ~ sP27, inference(subsumption_resolution, [], [f423, f168])).
fof(f423, plain, (~ (e1 = op(e0, e3)) | ~ sP27), inference(cnf_transformation, [], [f127])).
fof(f127, plain, ((~ (e1 = op(e0, e4)) & ~ (e1 = op(e0, e3)) & ~ (e1 = op(e0, e2)) & ~ (e1 = op(e0, e1)) & ~ (e1 = op(e0, e0))) | ~ sP27), inference(nnf_transformation, [], [f33])).
fof(f33, plain, ((~ (e1 = op(e0, e4)) & ~ (e1 = op(e0, e3)) & ~ (e1 = op(e0, e2)) & ~ (e1 = op(e0, e1)) & ~ (e1 = op(e0, e0))) | ~ sP27), inference(usedef, [], [e33])).
fof(e33, plain, (sP27 <=> (~ (e1 = op(e0, e4)) & ~ (e1 = op(e0, e3)) & ~ (e1 = op(e0, e2)) & ~ (e1 = op(e0, e1)) & ~ (e1 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP27])])).
fof(f835, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f834, f701])).
fof(f701, plain, ~ sP28, inference(subsumption_resolution, [], [f417, f175])).
fof(f417, plain, (~ (e1 = op(e2, e0)) | ~ sP28), inference(cnf_transformation, [], [f126])).
fof(f126, plain, ((~ (e1 = op(e4, e0)) & ~ (e1 = op(e3, e0)) & ~ (e1 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e1 = op(e0, e0))) | ~ sP28), inference(nnf_transformation, [], [f34])).
fof(f34, plain, ((~ (e1 = op(e4, e0)) & ~ (e1 = op(e3, e0)) & ~ (e1 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e1 = op(e0, e0))) | ~ sP28), inference(usedef, [], [e34])).
fof(e34, plain, (sP28 <=> (~ (e1 = op(e4, e0)) & ~ (e1 = op(e3, e0)) & ~ (e1 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e1 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP28])])).
fof(f834, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f833, f698])).
fof(f698, plain, ~ sP29, inference(subsumption_resolution, [], [f414, f169])).
fof(f414, plain, (~ (e2 = op(e0, e4)) | ~ sP29), inference(cnf_transformation, [], [f125])).
fof(f125, plain, ((~ (e2 = op(e0, e4)) & ~ (e2 = op(e0, e3)) & ~ (e2 = op(e0, e2)) & ~ (e2 = op(e0, e1)) & ~ (e2 = op(e0, e0))) | ~ sP29), inference(nnf_transformation, [], [f35])).
fof(f35, plain, ((~ (e2 = op(e0, e4)) & ~ (e2 = op(e0, e3)) & ~ (e2 = op(e0, e2)) & ~ (e2 = op(e0, e1)) & ~ (e2 = op(e0, e0))) | ~ sP29), inference(usedef, [], [e35])).
fof(e35, plain, (sP29 <=> (~ (e2 = op(e0, e4)) & ~ (e2 = op(e0, e3)) & ~ (e2 = op(e0, e2)) & ~ (e2 = op(e0, e1)) & ~ (e2 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP29])])).
fof(f833, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f832, f697])).
fof(f697, plain, ~ sP30, inference(subsumption_resolution, [], [f408, f180])).
fof(f408, plain, (~ (e2 = op(e3, e0)) | ~ sP30), inference(cnf_transformation, [], [f124])).
fof(f124, plain, ((~ (e2 = op(e4, e0)) & ~ (e2 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e2 = op(e1, e0)) & ~ (e2 = op(e0, e0))) | ~ sP30), inference(nnf_transformation, [], [f36])).
fof(f36, plain, ((~ (e2 = op(e4, e0)) & ~ (e2 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e2 = op(e1, e0)) & ~ (e2 = op(e0, e0))) | ~ sP30), inference(usedef, [], [e36])).
fof(e36, plain, (sP30 <=> (~ (e2 = op(e4, e0)) & ~ (e2 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e2 = op(e1, e0)) & ~ (e2 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP30])])).
fof(f832, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f831, f695])).
fof(f695, plain, ~ sP31, inference(subsumption_resolution, [], [f401, f166])).
fof(f401, plain, (~ (e3 = op(e0, e1)) | ~ sP31), inference(cnf_transformation, [], [f123])).
fof(f123, plain, ((~ (e3 = op(e0, e4)) & ~ (e3 = op(e0, e3)) & ~ (e3 = op(e0, e2)) & ~ (e3 = op(e0, e1)) & ~ (e3 = op(e0, e0))) | ~ sP31), inference(nnf_transformation, [], [f37])).
fof(f37, plain, ((~ (e3 = op(e0, e4)) & ~ (e3 = op(e0, e3)) & ~ (e3 = op(e0, e2)) & ~ (e3 = op(e0, e1)) & ~ (e3 = op(e0, e0))) | ~ sP31), inference(usedef, [], [e37])).
fof(e37, plain, (sP31 <=> (~ (e3 = op(e0, e4)) & ~ (e3 = op(e0, e3)) & ~ (e3 = op(e0, e2)) & ~ (e3 = op(e0, e1)) & ~ (e3 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP31])])).
fof(f831, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f830, f691])).
fof(f691, plain, ~ sP32, inference(subsumption_resolution, [], [f399, f185])).
fof(f399, plain, (~ (e3 = op(e4, e0)) | ~ sP32), inference(cnf_transformation, [], [f122])).
fof(f122, plain, ((~ (e3 = op(e4, e0)) & ~ (e3 = op(e3, e0)) & ~ (e3 = op(e2, e0)) & ~ (e3 = op(e1, e0)) & ~ (e3 = op(e0, e0))) | ~ sP32), inference(nnf_transformation, [], [f38])).
fof(f38, plain, ((~ (e3 = op(e4, e0)) & ~ (e3 = op(e3, e0)) & ~ (e3 = op(e2, e0)) & ~ (e3 = op(e1, e0)) & ~ (e3 = op(e0, e0))) | ~ sP32), inference(usedef, [], [e38])).
fof(e38, plain, (sP32 <=> (~ (e3 = op(e4, e0)) & ~ (e3 = op(e3, e0)) & ~ (e3 = op(e2, e0)) & ~ (e3 = op(e1, e0)) & ~ (e3 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP32])])).
fof(f830, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f829, f690])).
fof(f690, plain, ~ sP33, inference(subsumption_resolution, [], [f392, f167])).
fof(f392, plain, (~ (e4 = op(e0, e2)) | ~ sP33), inference(cnf_transformation, [], [f121])).
fof(f121, plain, ((~ (e4 = op(e0, e4)) & ~ (e4 = op(e0, e3)) & ~ (e4 = op(e0, e2)) & ~ (e4 = op(e0, e1)) & ~ (e4 = op(e0, e0))) | ~ sP33), inference(nnf_transformation, [], [f39])).
fof(f39, plain, ((~ (e4 = op(e0, e4)) & ~ (e4 = op(e0, e3)) & ~ (e4 = op(e0, e2)) & ~ (e4 = op(e0, e1)) & ~ (e4 = op(e0, e0))) | ~ sP33), inference(usedef, [], [e39])).
fof(e39, plain, (sP33 <=> (~ (e4 = op(e0, e4)) & ~ (e4 = op(e0, e3)) & ~ (e4 = op(e0, e2)) & ~ (e4 = op(e0, e1)) & ~ (e4 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP33])])).
fof(f829, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f828, f687])).
fof(f687, plain, ~ sP34, inference(subsumption_resolution, [], [f386, f170])).
fof(f386, plain, (~ (e4 = op(e1, e0)) | ~ sP34), inference(cnf_transformation, [], [f120])).
fof(f120, plain, ((~ (e4 = op(e4, e0)) & ~ (e4 = op(e3, e0)) & ~ (e4 = op(e2, e0)) & ~ (e4 = op(e1, e0)) & ~ (e4 = op(e0, e0))) | ~ sP34), inference(nnf_transformation, [], [f40])).
fof(f40, plain, ((~ (e4 = op(e4, e0)) & ~ (e4 = op(e3, e0)) & ~ (e4 = op(e2, e0)) & ~ (e4 = op(e1, e0)) & ~ (e4 = op(e0, e0))) | ~ sP34), inference(usedef, [], [e40])).
fof(e40, plain, (sP34 <=> (~ (e4 = op(e4, e0)) & ~ (e4 = op(e3, e0)) & ~ (e4 = op(e2, e0)) & ~ (e4 = op(e1, e0)) & ~ (e4 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP34])])).
fof(f828, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f827, f683])).
fof(f683, plain, ~ sP35, inference(subsumption_resolution, [], [f384, f174])).
fof(f384, plain, (~ (e0 = op(e1, e4)) | ~ sP35), inference(cnf_transformation, [], [f119])).
fof(f119, plain, ((~ (e0 = op(e1, e4)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e1, e0))) | ~ sP35), inference(nnf_transformation, [], [f41])).
fof(f41, plain, ((~ (e0 = op(e1, e4)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e1, e0))) | ~ sP35), inference(usedef, [], [e41])).
fof(e41, plain, (sP35 <=> (~ (e0 = op(e1, e4)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP35])])).
fof(f827, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f826, f682])).
fof(f682, plain, ~ sP36, inference(subsumption_resolution, [], [f377, f176])).
fof(f377, plain, (~ (e0 = op(e2, e1)) | ~ sP36), inference(cnf_transformation, [], [f118])).
fof(f118, plain, ((~ (e0 = op(e4, e1)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e0, e1))) | ~ sP36), inference(nnf_transformation, [], [f42])).
fof(f42, plain, ((~ (e0 = op(e4, e1)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e0, e1))) | ~ sP36), inference(usedef, [], [e42])).
fof(e42, plain, (sP36 <=> (~ (e0 = op(e4, e1)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP36])])).
fof(f826, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f825, f679])).
fof(f679, plain, ~ sP37, inference(subsumption_resolution, [], [f371, f171])).
fof(f371, plain, (~ (e1 = op(e1, e1)) | ~ sP37), inference(cnf_transformation, [], [f117])).
fof(f117, plain, ((~ (e1 = op(e1, e4)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e1, e0))) | ~ sP37), inference(nnf_transformation, [], [f43])).
fof(f43, plain, ((~ (e1 = op(e1, e4)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e1, e0))) | ~ sP37), inference(usedef, [], [e43])).
fof(e43, plain, (sP37 <=> (~ (e1 = op(e1, e4)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP37])])).
fof(f825, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f824, f675])).
fof(f675, plain, ~ sP38, inference(subsumption_resolution, [], [f366, f171])).
fof(f366, plain, (~ (e1 = op(e1, e1)) | ~ sP38), inference(cnf_transformation, [], [f116])).
fof(f116, plain, ((~ (e1 = op(e4, e1)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e0, e1))) | ~ sP38), inference(nnf_transformation, [], [f44])).
fof(f44, plain, ((~ (e1 = op(e4, e1)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e0, e1))) | ~ sP38), inference(usedef, [], [e44])).
fof(e44, plain, (sP38 <=> (~ (e1 = op(e4, e1)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP38])])).
fof(f824, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f823, f671])).
fof(f671, plain, ~ sP39, inference(subsumption_resolution, [], [f363, f173])).
fof(f363, plain, (~ (e2 = op(e1, e3)) | ~ sP39), inference(cnf_transformation, [], [f115])).
fof(f115, plain, ((~ (e2 = op(e1, e4)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e1, e0))) | ~ sP39), inference(nnf_transformation, [], [f45])).
fof(f45, plain, ((~ (e2 = op(e1, e4)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e1, e0))) | ~ sP39), inference(usedef, [], [e45])).
fof(e45, plain, (sP39 <=> (~ (e2 = op(e1, e4)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP39])])).
fof(f823, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f822, f669])).
fof(f669, plain, ~ sP40, inference(subsumption_resolution, [], [f359, f186])).
fof(f359, plain, (~ (e2 = op(e4, e1)) | ~ sP40), inference(cnf_transformation, [], [f114])).
fof(f114, plain, ((~ (e2 = op(e4, e1)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e0, e1))) | ~ sP40), inference(nnf_transformation, [], [f46])).
fof(f46, plain, ((~ (e2 = op(e4, e1)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e0, e1))) | ~ sP40), inference(usedef, [], [e46])).
fof(e46, plain, (sP40 <=> (~ (e2 = op(e4, e1)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP40])])).
fof(f822, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f821, f668])).
fof(f668, plain, ~ sP41, inference(subsumption_resolution, [], [f352, f172])).
fof(f352, plain, (~ (e3 = op(e1, e2)) | ~ sP41), inference(cnf_transformation, [], [f113])).
fof(f113, plain, ((~ (e3 = op(e1, e4)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e1, e0))) | ~ sP41), inference(nnf_transformation, [], [f47])).
fof(f47, plain, ((~ (e3 = op(e1, e4)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e1, e0))) | ~ sP41), inference(usedef, [], [e47])).
fof(e47, plain, (sP41 <=> (~ (e3 = op(e1, e4)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP41])])).
fof(f821, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f820, f665])).
fof(f665, plain, ~ sP42, inference(subsumption_resolution, [], [f345, f166])).
fof(f345, plain, (~ (e3 = op(e0, e1)) | ~ sP42), inference(cnf_transformation, [], [f112])).
fof(f112, plain, ((~ (e3 = op(e4, e1)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e0, e1))) | ~ sP42), inference(nnf_transformation, [], [f48])).
fof(f48, plain, ((~ (e3 = op(e4, e1)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e0, e1))) | ~ sP42), inference(usedef, [], [e48])).
fof(e48, plain, (sP42 <=> (~ (e3 = op(e4, e1)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP42])])).
fof(f820, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f819, f660])).
fof(f660, plain, ~ sP43, inference(subsumption_resolution, [], [f340, f170])).
fof(f340, plain, (~ (e4 = op(e1, e0)) | ~ sP43), inference(cnf_transformation, [], [f111])).
fof(f111, plain, ((~ (e4 = op(e1, e4)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e1, e0))) | ~ sP43), inference(nnf_transformation, [], [f49])).
fof(f49, plain, ((~ (e4 = op(e1, e4)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e1, e0))) | ~ sP43), inference(usedef, [], [e49])).
fof(e49, plain, (sP43 <=> (~ (e4 = op(e1, e4)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP43])])).
fof(f819, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f818, f655])).
fof(f655, plain, ~ sP44, inference(subsumption_resolution, [], [f338, f181])).
fof(f338, plain, (~ (e4 = op(e3, e1)) | ~ sP44), inference(cnf_transformation, [], [f110])).
fof(f110, plain, ((~ (e4 = op(e4, e1)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e0, e1))) | ~ sP44), inference(nnf_transformation, [], [f50])).
fof(f50, plain, ((~ (e4 = op(e4, e1)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e0, e1))) | ~ sP44), inference(usedef, [], [e50])).
fof(e50, plain, (sP44 <=> (~ (e4 = op(e4, e1)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP44])])).
fof(f818, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f817, f653])).
fof(f653, plain, ~ sP45, inference(subsumption_resolution, [], [f331, f176])).
fof(f331, plain, (~ (e0 = op(e2, e1)) | ~ sP45), inference(cnf_transformation, [], [f109])).
fof(f109, plain, ((~ (e0 = op(e2, e4)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e2, e0))) | ~ sP45), inference(nnf_transformation, [], [f51])).
fof(f51, plain, ((~ (e0 = op(e2, e4)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e2, e0))) | ~ sP45), inference(usedef, [], [e51])).
fof(e51, plain, (sP45 <=> (~ (e0 = op(e2, e4)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP45])])).
fof(f817, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f816, f649])).
fof(f649, plain, ~ sP46, inference(subsumption_resolution, [], [f328, f182])).
fof(f328, plain, (~ (e0 = op(e3, e2)) | ~ sP46), inference(cnf_transformation, [], [f108])).
fof(f108, plain, ((~ (e0 = op(e4, e2)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e0, e2))) | ~ sP46), inference(nnf_transformation, [], [f52])).
fof(f52, plain, ((~ (e0 = op(e4, e2)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e0, e2))) | ~ sP46), inference(usedef, [], [e52])).
fof(e52, plain, (sP46 <=> (~ (e0 = op(e4, e2)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP46])])).
fof(f816, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f815, f647])).
fof(f647, plain, ~ sP47, inference(subsumption_resolution, [], [f320, f175])).
fof(f320, plain, (~ (e1 = op(e2, e0)) | ~ sP47), inference(cnf_transformation, [], [f107])).
fof(f107, plain, ((~ (e1 = op(e2, e4)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e2, e0))) | ~ sP47), inference(nnf_transformation, [], [f53])).
fof(f53, plain, ((~ (e1 = op(e2, e4)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e2, e0))) | ~ sP47), inference(usedef, [], [e53])).
fof(e53, plain, (sP47 <=> (~ (e1 = op(e2, e4)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP47])])).
fof(f815, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f814, f642])).
fof(f642, plain, ~ sP48, inference(subsumption_resolution, [], [f319, f187])).
fof(f319, plain, (~ (e1 = op(e4, e2)) | ~ sP48), inference(cnf_transformation, [], [f106])).
fof(f106, plain, ((~ (e1 = op(e4, e2)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e0, e2))) | ~ sP48), inference(nnf_transformation, [], [f54])).
fof(f54, plain, ((~ (e1 = op(e4, e2)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e0, e2))) | ~ sP48), inference(usedef, [], [e54])).
fof(e54, plain, (sP48 <=> (~ (e1 = op(e4, e2)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP48])])).
fof(f814, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f813, f641])).
fof(f641, plain, ~ sP49, inference(subsumption_resolution, [], [f312, f177])).
fof(f312, plain, (~ (e2 = op(e2, e2)) | ~ sP49), inference(cnf_transformation, [], [f105])).
fof(f105, plain, ((~ (e2 = op(e2, e4)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e2, e0))) | ~ sP49), inference(nnf_transformation, [], [f55])).
fof(f55, plain, ((~ (e2 = op(e2, e4)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e2, e0))) | ~ sP49), inference(usedef, [], [e55])).
fof(e55, plain, (sP49 <=> (~ (e2 = op(e2, e4)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP49])])).
fof(f813, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f812, f638])).
fof(f638, plain, ~ sP50, inference(subsumption_resolution, [], [f307, f177])).
fof(f307, plain, (~ (e2 = op(e2, e2)) | ~ sP50), inference(cnf_transformation, [], [f104])).
fof(f104, plain, ((~ (e2 = op(e4, e2)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e0, e2))) | ~ sP50), inference(nnf_transformation, [], [f56])).
fof(f56, plain, ((~ (e2 = op(e4, e2)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e0, e2))) | ~ sP50), inference(usedef, [], [e56])).
fof(e56, plain, (sP50 <=> (~ (e2 = op(e4, e2)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP50])])).
fof(f812, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f811, f635])).
fof(f635, plain, ~ sP51, inference(subsumption_resolution, [], [f304, f179])).
fof(f304, plain, (~ (e3 = op(e2, e4)) | ~ sP51), inference(cnf_transformation, [], [f103])).
fof(f103, plain, ((~ (e3 = op(e2, e4)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e2, e0))) | ~ sP51), inference(nnf_transformation, [], [f57])).
fof(f57, plain, ((~ (e3 = op(e2, e4)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e2, e0))) | ~ sP51), inference(usedef, [], [e57])).
fof(e57, plain, (sP51 <=> (~ (e3 = op(e2, e4)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP51])])).
fof(f811, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f810, f634])).
fof(f634, plain, ~ sP52, inference(subsumption_resolution, [], [f296, f172])).
fof(f296, plain, (~ (e3 = op(e1, e2)) | ~ sP52), inference(cnf_transformation, [], [f102])).
fof(f102, plain, ((~ (e3 = op(e4, e2)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e0, e2))) | ~ sP52), inference(nnf_transformation, [], [f58])).
fof(f58, plain, ((~ (e3 = op(e4, e2)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e0, e2))) | ~ sP52), inference(usedef, [], [e58])).
fof(e58, plain, (sP52 <=> (~ (e3 = op(e4, e2)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP52])])).
fof(f810, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f809, f630])).
fof(f630, plain, ~ sP53, inference(subsumption_resolution, [], [f293, f178])).
fof(f293, plain, (~ (e4 = op(e2, e3)) | ~ sP53), inference(cnf_transformation, [], [f101])).
fof(f101, plain, ((~ (e4 = op(e2, e4)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e2, e0))) | ~ sP53), inference(nnf_transformation, [], [f59])).
fof(f59, plain, ((~ (e4 = op(e2, e4)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e2, e0))) | ~ sP53), inference(usedef, [], [e59])).
fof(e59, plain, (sP53 <=> (~ (e4 = op(e2, e4)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP53])])).
fof(f809, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f808, f628])).
fof(f628, plain, ~ sP54, inference(subsumption_resolution, [], [f285, f167])).
fof(f285, plain, (~ (e4 = op(e0, e2)) | ~ sP54), inference(cnf_transformation, [], [f100])).
fof(f100, plain, ((~ (e4 = op(e4, e2)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e0, e2))) | ~ sP54), inference(nnf_transformation, [], [f60])).
fof(f60, plain, ((~ (e4 = op(e4, e2)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e0, e2))) | ~ sP54), inference(usedef, [], [e60])).
fof(e60, plain, (sP54 <=> (~ (e4 = op(e4, e2)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP54])])).
fof(f808, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f807, f623])).
fof(f623, plain, ~ sP55, inference(subsumption_resolution, [], [f282, f182])).
fof(f282, plain, (~ (e0 = op(e3, e2)) | ~ sP55), inference(cnf_transformation, [], [f99])).
fof(f99, plain, ((~ (e0 = op(e3, e4)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e3, e0))) | ~ sP55), inference(nnf_transformation, [], [f61])).
fof(f61, plain, ((~ (e0 = op(e3, e4)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e3, e0))) | ~ sP55), inference(usedef, [], [e61])).
fof(e61, plain, (sP55 <=> (~ (e0 = op(e3, e4)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP55])])).
fof(f807, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f806, f620])).
fof(f620, plain, ~ sP56, inference(subsumption_resolution, [], [f279, f188])).
fof(f279, plain, (~ (e0 = op(e4, e3)) | ~ sP56), inference(cnf_transformation, [], [f98])).
fof(f98, plain, ((~ (e0 = op(e4, e3)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e0, e3))) | ~ sP56), inference(nnf_transformation, [], [f62])).
fof(f62, plain, ((~ (e0 = op(e4, e3)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e0, e3))) | ~ sP56), inference(usedef, [], [e62])).
fof(e62, plain, (sP56 <=> (~ (e0 = op(e4, e3)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP56])])).
fof(f806, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f805, f619])).
fof(f619, plain, ~ sP57, inference(subsumption_resolution, [], [f274, f184])).
fof(f274, plain, (~ (e1 = op(e3, e4)) | ~ sP57), inference(cnf_transformation, [], [f97])).
fof(f97, plain, ((~ (e1 = op(e3, e4)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e3, e0))) | ~ sP57), inference(nnf_transformation, [], [f63])).
fof(f63, plain, ((~ (e1 = op(e3, e4)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e3, e0))) | ~ sP57), inference(usedef, [], [e63])).
fof(e63, plain, (sP57 <=> (~ (e1 = op(e3, e4)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP57])])).
fof(f805, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f804, f618])).
fof(f618, plain, ~ sP58, inference(subsumption_resolution, [], [f265, f168])).
fof(f265, plain, (~ (e1 = op(e0, e3)) | ~ sP58), inference(cnf_transformation, [], [f96])).
fof(f96, plain, ((~ (e1 = op(e4, e3)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e0, e3))) | ~ sP58), inference(nnf_transformation, [], [f64])).
fof(f64, plain, ((~ (e1 = op(e4, e3)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e0, e3))) | ~ sP58), inference(usedef, [], [e64])).
fof(e64, plain, (sP58 <=> (~ (e1 = op(e4, e3)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP58])])).
fof(f804, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f803, f613])).
fof(f613, plain, ~ sP59, inference(subsumption_resolution, [], [f260, f180])).
fof(f260, plain, (~ (e2 = op(e3, e0)) | ~ sP59), inference(cnf_transformation, [], [f95])).
fof(f95, plain, ((~ (e2 = op(e3, e4)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e3, e0))) | ~ sP59), inference(nnf_transformation, [], [f65])).
fof(f65, plain, ((~ (e2 = op(e3, e4)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e3, e0))) | ~ sP59), inference(usedef, [], [e65])).
fof(e65, plain, (sP59 <=> (~ (e2 = op(e3, e4)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP59])])).
fof(f803, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f802, f608])).
fof(f608, plain, ~ sP60, inference(subsumption_resolution, [], [f256, f173])).
fof(f256, plain, (~ (e2 = op(e1, e3)) | ~ sP60), inference(cnf_transformation, [], [f94])).
fof(f94, plain, ((~ (e2 = op(e4, e3)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e0, e3))) | ~ sP60), inference(nnf_transformation, [], [f66])).
fof(f66, plain, ((~ (e2 = op(e4, e3)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e0, e3))) | ~ sP60), inference(usedef, [], [e66])).
fof(e66, plain, (sP60 <=> (~ (e2 = op(e4, e3)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP60])])).
fof(f802, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f801, f604])).
fof(f604, plain, ~ sP61, inference(subsumption_resolution, [], [f253, f183])).
fof(f253, plain, (~ (e3 = op(e3, e3)) | ~ sP61), inference(cnf_transformation, [], [f93])).
fof(f93, plain, ((~ (e3 = op(e3, e4)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e3, e0))) | ~ sP61), inference(nnf_transformation, [], [f67])).
fof(f67, plain, ((~ (e3 = op(e3, e4)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e3, e0))) | ~ sP61), inference(usedef, [], [e67])).
fof(e67, plain, (sP61 <=> (~ (e3 = op(e3, e4)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP61])])).
fof(f801, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f800, f602])).
fof(f602, plain, ~ sP62, inference(subsumption_resolution, [], [f248, f183])).
fof(f248, plain, (~ (e3 = op(e3, e3)) | ~ sP62), inference(cnf_transformation, [], [f92])).
fof(f92, plain, ((~ (e3 = op(e4, e3)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e0, e3))) | ~ sP62), inference(nnf_transformation, [], [f68])).
fof(f68, plain, ((~ (e3 = op(e4, e3)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e0, e3))) | ~ sP62), inference(usedef, [], [e68])).
fof(e68, plain, (sP62 <=> (~ (e3 = op(e4, e3)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP62])])).
fof(f800, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f799, f600])).
fof(f600, plain, ~ sP63, inference(subsumption_resolution, [], [f241, f181])).
fof(f241, plain, (~ (e4 = op(e3, e1)) | ~ sP63), inference(cnf_transformation, [], [f91])).
fof(f91, plain, ((~ (e4 = op(e3, e4)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e3, e0))) | ~ sP63), inference(nnf_transformation, [], [f69])).
fof(f69, plain, ((~ (e4 = op(e3, e4)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e3, e0))) | ~ sP63), inference(usedef, [], [e69])).
fof(e69, plain, (sP63 <=> (~ (e4 = op(e3, e4)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP63])])).
fof(f799, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f798, f596])).
fof(f596, plain, ~ sP64, inference(subsumption_resolution, [], [f237, f178])).
fof(f237, plain, (~ (e4 = op(e2, e3)) | ~ sP64), inference(cnf_transformation, [], [f90])).
fof(f90, plain, ((~ (e4 = op(e4, e3)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e0, e3))) | ~ sP64), inference(nnf_transformation, [], [f70])).
fof(f70, plain, ((~ (e4 = op(e4, e3)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e0, e3))) | ~ sP64), inference(usedef, [], [e70])).
fof(e70, plain, (sP64 <=> (~ (e4 = op(e4, e3)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP64])])).
fof(f798, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f797, f593])).
fof(f593, plain, ~ sP65, inference(subsumption_resolution, [], [f233, f188])).
fof(f233, plain, (~ (e0 = op(e4, e3)) | ~ sP65), inference(cnf_transformation, [], [f89])).
fof(f89, plain, ((~ (e0 = op(e4, e4)) & ~ (e0 = op(e4, e3)) & ~ (e0 = op(e4, e2)) & ~ (e0 = op(e4, e1)) & ~ (e0 = op(e4, e0))) | ~ sP65), inference(nnf_transformation, [], [f71])).
fof(f71, plain, ((~ (e0 = op(e4, e4)) & ~ (e0 = op(e4, e3)) & ~ (e0 = op(e4, e2)) & ~ (e0 = op(e4, e1)) & ~ (e0 = op(e4, e0))) | ~ sP65), inference(usedef, [], [e71])).
fof(e71, plain, (sP65 <=> (~ (e0 = op(e4, e4)) & ~ (e0 = op(e4, e3)) & ~ (e0 = op(e4, e2)) & ~ (e0 = op(e4, e1)) & ~ (e0 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP65])])).
fof(f797, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f796, f591])).
fof(f591, plain, ~ sP66, inference(subsumption_resolution, [], [f226, f174])).
fof(f226, plain, (~ (e0 = op(e1, e4)) | ~ sP66), inference(cnf_transformation, [], [f88])).
fof(f88, plain, ((~ (e0 = op(e4, e4)) & ~ (e0 = op(e3, e4)) & ~ (e0 = op(e2, e4)) & ~ (e0 = op(e1, e4)) & ~ (e0 = op(e0, e4))) | ~ sP66), inference(nnf_transformation, [], [f72])).
fof(f72, plain, ((~ (e0 = op(e4, e4)) & ~ (e0 = op(e3, e4)) & ~ (e0 = op(e2, e4)) & ~ (e0 = op(e1, e4)) & ~ (e0 = op(e0, e4))) | ~ sP66), inference(usedef, [], [e72])).
fof(e72, plain, (sP66 <=> (~ (e0 = op(e4, e4)) & ~ (e0 = op(e3, e4)) & ~ (e0 = op(e2, e4)) & ~ (e0 = op(e1, e4)) & ~ (e0 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP66])])).
fof(f796, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f795, f587])).
fof(f587, plain, ~ sP67, inference(subsumption_resolution, [], [f222, f187])).
fof(f222, plain, (~ (e1 = op(e4, e2)) | ~ sP67), inference(cnf_transformation, [], [f87])).
fof(f87, plain, ((~ (e1 = op(e4, e4)) & ~ (e1 = op(e4, e3)) & ~ (e1 = op(e4, e2)) & ~ (e1 = op(e4, e1)) & ~ (e1 = op(e4, e0))) | ~ sP67), inference(nnf_transformation, [], [f73])).
fof(f73, plain, ((~ (e1 = op(e4, e4)) & ~ (e1 = op(e4, e3)) & ~ (e1 = op(e4, e2)) & ~ (e1 = op(e4, e1)) & ~ (e1 = op(e4, e0))) | ~ sP67), inference(usedef, [], [e73])).
fof(e73, plain, (sP67 <=> (~ (e1 = op(e4, e4)) & ~ (e1 = op(e4, e3)) & ~ (e1 = op(e4, e2)) & ~ (e1 = op(e4, e1)) & ~ (e1 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP67])])).
fof(f795, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f794, f584])).
fof(f584, plain, ~ sP68, inference(subsumption_resolution, [], [f218, f184])).
fof(f218, plain, (~ (e1 = op(e3, e4)) | ~ sP68), inference(cnf_transformation, [], [f86])).
fof(f86, plain, ((~ (e1 = op(e4, e4)) & ~ (e1 = op(e3, e4)) & ~ (e1 = op(e2, e4)) & ~ (e1 = op(e1, e4)) & ~ (e1 = op(e0, e4))) | ~ sP68), inference(nnf_transformation, [], [f74])).
fof(f74, plain, ((~ (e1 = op(e4, e4)) & ~ (e1 = op(e3, e4)) & ~ (e1 = op(e2, e4)) & ~ (e1 = op(e1, e4)) & ~ (e1 = op(e0, e4))) | ~ sP68), inference(usedef, [], [e74])).
fof(e74, plain, (sP68 <=> (~ (e1 = op(e4, e4)) & ~ (e1 = op(e3, e4)) & ~ (e1 = op(e2, e4)) & ~ (e1 = op(e1, e4)) & ~ (e1 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP68])])).
fof(f794, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f793, f582])).
fof(f582, plain, ~ sP69, inference(subsumption_resolution, [], [f211, f186])).
fof(f211, plain, (~ (e2 = op(e4, e1)) | ~ sP69), inference(cnf_transformation, [], [f85])).
fof(f85, plain, ((~ (e2 = op(e4, e4)) & ~ (e2 = op(e4, e3)) & ~ (e2 = op(e4, e2)) & ~ (e2 = op(e4, e1)) & ~ (e2 = op(e4, e0))) | ~ sP69), inference(nnf_transformation, [], [f75])).
fof(f75, plain, ((~ (e2 = op(e4, e4)) & ~ (e2 = op(e4, e3)) & ~ (e2 = op(e4, e2)) & ~ (e2 = op(e4, e1)) & ~ (e2 = op(e4, e0))) | ~ sP69), inference(usedef, [], [e75])).
fof(e75, plain, (sP69 <=> (~ (e2 = op(e4, e4)) & ~ (e2 = op(e4, e3)) & ~ (e2 = op(e4, e2)) & ~ (e2 = op(e4, e1)) & ~ (e2 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP69])])).
fof(f793, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f792, f578])).
fof(f578, plain, ~ sP70, inference(subsumption_resolution, [], [f205, f169])).
fof(f205, plain, (~ (e2 = op(e0, e4)) | ~ sP70), inference(cnf_transformation, [], [f84])).
fof(f84, plain, ((~ (e2 = op(e4, e4)) & ~ (e2 = op(e3, e4)) & ~ (e2 = op(e2, e4)) & ~ (e2 = op(e1, e4)) & ~ (e2 = op(e0, e4))) | ~ sP70), inference(nnf_transformation, [], [f76])).
fof(f76, plain, ((~ (e2 = op(e4, e4)) & ~ (e2 = op(e3, e4)) & ~ (e2 = op(e2, e4)) & ~ (e2 = op(e1, e4)) & ~ (e2 = op(e0, e4))) | ~ sP70), inference(usedef, [], [e76])).
fof(e76, plain, (sP70 <=> (~ (e2 = op(e4, e4)) & ~ (e2 = op(e3, e4)) & ~ (e2 = op(e2, e4)) & ~ (e2 = op(e1, e4)) & ~ (e2 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP70])])).
fof(f792, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f791, f573])).
fof(f573, plain, ~ sP71, inference(subsumption_resolution, [], [f200, f185])).
fof(f200, plain, (~ (e3 = op(e4, e0)) | ~ sP71), inference(cnf_transformation, [], [f83])).
fof(f83, plain, ((~ (e3 = op(e4, e4)) & ~ (e3 = op(e4, e3)) & ~ (e3 = op(e4, e2)) & ~ (e3 = op(e4, e1)) & ~ (e3 = op(e4, e0))) | ~ sP71), inference(nnf_transformation, [], [f77])).
fof(f77, plain, ((~ (e3 = op(e4, e4)) & ~ (e3 = op(e4, e3)) & ~ (e3 = op(e4, e2)) & ~ (e3 = op(e4, e1)) & ~ (e3 = op(e4, e0))) | ~ sP71), inference(usedef, [], [e77])).
fof(e77, plain, (sP71 <=> (~ (e3 = op(e4, e4)) & ~ (e3 = op(e4, e3)) & ~ (e3 = op(e4, e2)) & ~ (e3 = op(e4, e1)) & ~ (e3 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP71])])).
fof(f791, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f790, f568])).
fof(f568, plain, ~ sP72, inference(subsumption_resolution, [], [f197, f179])).
fof(f197, plain, (~ (e3 = op(e2, e4)) | ~ sP72), inference(cnf_transformation, [], [f82])).
fof(f82, plain, ((~ (e3 = op(e4, e4)) & ~ (e3 = op(e3, e4)) & ~ (e3 = op(e2, e4)) & ~ (e3 = op(e1, e4)) & ~ (e3 = op(e0, e4))) | ~ sP72), inference(nnf_transformation, [], [f78])).
fof(f78, plain, ((~ (e3 = op(e4, e4)) & ~ (e3 = op(e3, e4)) & ~ (e3 = op(e2, e4)) & ~ (e3 = op(e1, e4)) & ~ (e3 = op(e0, e4))) | ~ sP72), inference(usedef, [], [e78])).
fof(e78, plain, (sP72 <=> (~ (e3 = op(e4, e4)) & ~ (e3 = op(e3, e4)) & ~ (e3 = op(e2, e4)) & ~ (e3 = op(e1, e4)) & ~ (e3 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP72])])).
fof(f790, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f789, f565])).
fof(f565, plain, ~ sP73, inference(subsumption_resolution, [], [f194, f189])).
fof(f194, plain, (~ (e4 = op(e4, e4)) | ~ sP73), inference(cnf_transformation, [], [f81])).
fof(f81, plain, ((~ (e4 = op(e4, e4)) & ~ (e4 = op(e4, e3)) & ~ (e4 = op(e4, e2)) & ~ (e4 = op(e4, e1)) & ~ (e4 = op(e4, e0))) | ~ sP73), inference(nnf_transformation, [], [f79])).
fof(f79, plain, ((~ (e4 = op(e4, e4)) & ~ (e4 = op(e4, e3)) & ~ (e4 = op(e4, e2)) & ~ (e4 = op(e4, e1)) & ~ (e4 = op(e4, e0))) | ~ sP73), inference(usedef, [], [e79])).
fof(e79, plain, (sP73 <=> (~ (e4 = op(e4, e4)) & ~ (e4 = op(e4, e3)) & ~ (e4 = op(e4, e2)) & ~ (e4 = op(e4, e1)) & ~ (e4 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP73])])).
fof(f789, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(subsumption_resolution, [], [f564, f189])).
fof(f564, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e4, e4)) | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(cnf_transformation, [], [f80])).
fof(f80, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | (~ (e4 = op(e4, e4)) & ~ (e4 = op(e3, e4)) & ~ (e4 = op(e2, e4)) & ~ (e4 = op(e1, e4)) & ~ (e4 = op(e0, e4))) | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(definition_folding, [], [f5, e79, e78, e77, e76, e75, e74, e73, e72, e71, e70, e69, e68, e67, e66, e65, e64, e63, e62, e61, e60, e59, e58, e57, e56, e55, e54, e53, e52, e51, e50, e49, e48, e47, e46, e45, e44, e43, e42, e41, e40, e39, e38, e37, e36, e35, e34, e33, e32, e31, e30, e29, e28, e27, e26, e25, e24, e23, e22, e21, e20, e19, e18, e17, e16, e15, e14, e13, e12, e11, e10, e9, e8, e7, e6])).
fof(f5, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | (~ (e4 = op(e4, e4)) & ~ (e4 = op(e3, e4)) & ~ (e4 = op(e2, e4)) & ~ (e4 = op(e1, e4)) & ~ (e4 = op(e0, e4))) | (~ (e4 = op(e4, e4)) & ~ (e4 = op(e4, e3)) & ~ (e4 = op(e4, e2)) & ~ (e4 = op(e4, e1)) & ~ (e4 = op(e4, e0))) | (~ (e3 = op(e4, e4)) & ~ (e3 = op(e3, e4)) & ~ (e3 = op(e2, e4)) & ~ (e3 = op(e1, e4)) & ~ (e3 = op(e0, e4))) | (~ (e3 = op(e4, e4)) & ~ (e3 = op(e4, e3)) & ~ (e3 = op(e4, e2)) & ~ (e3 = op(e4, e1)) & ~ (e3 = op(e4, e0))) | (~ (e2 = op(e4, e4)) & ~ (e2 = op(e3, e4)) & ~ (e2 = op(e2, e4)) & ~ (e2 = op(e1, e4)) & ~ (e2 = op(e0, e4))) | (~ (e2 = op(e4, e4)) & ~ (e2 = op(e4, e3)) & ~ (e2 = op(e4, e2)) & ~ (e2 = op(e4, e1)) & ~ (e2 = op(e4, e0))) | (~ (e1 = op(e4, e4)) & ~ (e1 = op(e3, e4)) & ~ (e1 = op(e2, e4)) & ~ (e1 = op(e1, e4)) & ~ (e1 = op(e0, e4))) | (~ (e1 = op(e4, e4)) & ~ (e1 = op(e4, e3)) & ~ (e1 = op(e4, e2)) & ~ (e1 = op(e4, e1)) & ~ (e1 = op(e4, e0))) | (~ (e0 = op(e4, e4)) & ~ (e0 = op(e3, e4)) & ~ (e0 = op(e2, e4)) & ~ (e0 = op(e1, e4)) & ~ (e0 = op(e0, e4))) | (~ (e0 = op(e4, e4)) & ~ (e0 = op(e4, e3)) & ~ (e0 = op(e4, e2)) & ~ (e0 = op(e4, e1)) & ~ (e0 = op(e4, e0))) | (~ (e4 = op(e4, e3)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e0, e3))) | (~ (e4 = op(e3, e4)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e3, e0))) | (~ (e3 = op(e4, e3)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e0, e3))) | (~ (e3 = op(e3, e4)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e3, e0))) | (~ (e2 = op(e4, e3)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e0, e3))) | (~ (e2 = op(e3, e4)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e3, e0))) | (~ (e1 = op(e4, e3)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e0, e3))) | (~ (e1 = op(e3, e4)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e3, e0))) | (~ (e0 = op(e4, e3)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e0, e3))) | (~ (e0 = op(e3, e4)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e3, e0))) | (~ (e4 = op(e4, e2)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e0, e2))) | (~ (e4 = op(e2, e4)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e2, e0))) | (~ (e3 = op(e4, e2)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e0, e2))) | (~ (e3 = op(e2, e4)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e2, e0))) | (~ (e2 = op(e4, e2)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e0, e2))) | (~ (e2 = op(e2, e4)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e2, e0))) | (~ (e1 = op(e4, e2)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e0, e2))) | (~ (e1 = op(e2, e4)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e2, e0))) | (~ (e0 = op(e4, e2)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e0, e2))) | (~ (e0 = op(e2, e4)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e2, e0))) | (~ (e4 = op(e4, e1)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e0, e1))) | (~ (e4 = op(e1, e4)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e1, e0))) | (~ (e3 = op(e4, e1)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e0, e1))) | (~ (e3 = op(e1, e4)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e1, e0))) | (~ (e2 = op(e4, e1)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e0, e1))) | (~ (e2 = op(e1, e4)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e1, e0))) | (~ (e1 = op(e4, e1)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e0, e1))) | (~ (e1 = op(e1, e4)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e1, e0))) | (~ (e0 = op(e4, e1)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e0, e1))) | (~ (e0 = op(e1, e4)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e1, e0))) | (~ (e4 = op(e4, e0)) & ~ (e4 = op(e3, e0)) & ~ (e4 = op(e2, e0)) & ~ (e4 = op(e1, e0)) & ~ (e4 = op(e0, e0))) | (~ (e4 = op(e0, e4)) & ~ (e4 = op(e0, e3)) & ~ (e4 = op(e0, e2)) & ~ (e4 = op(e0, e1)) & ~ (e4 = op(e0, e0))) | (~ (e3 = op(e4, e0)) & ~ (e3 = op(e3, e0)) & ~ (e3 = op(e2, e0)) & ~ (e3 = op(e1, e0)) & ~ (e3 = op(e0, e0))) | (~ (e3 = op(e0, e4)) & ~ (e3 = op(e0, e3)) & ~ (e3 = op(e0, e2)) & ~ (e3 = op(e0, e1)) & ~ (e3 = op(e0, e0))) | (~ (e2 = op(e4, e0)) & ~ (e2 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e2 = op(e1, e0)) & ~ (e2 = op(e0, e0))) | (~ (e2 = op(e0, e4)) & ~ (e2 = op(e0, e3)) & ~ (e2 = op(e0, e2)) & ~ (e2 = op(e0, e1)) & ~ (e2 = op(e0, e0))) | (~ (e1 = op(e4, e0)) & ~ (e1 = op(e3, e0)) & ~ (e1 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e1 = op(e0, e0))) | (~ (e1 = op(e0, e4)) & ~ (e1 = op(e0, e3)) & ~ (e1 = op(e0, e2)) & ~ (e1 = op(e0, e1)) & ~ (e1 = op(e0, e0))) | (~ (e0 = op(e4, e0)) & ~ (e0 = op(e3, e0)) & ~ (e0 = op(e2, e0)) & ~ (e0 = op(e1, e0)) & ~ (e0 = op(e0, e0))) | (~ (e0 = op(e0, e4)) & ~ (e0 = op(e0, e3)) & ~ (e0 = op(e0, e2)) & ~ (e0 = op(e0, e1)) & ~ (e0 = op(e0, e0))) | (~ (e4 = op(e4, e4)) & ~ (e3 = op(e4, e4)) & ~ (e2 = op(e4, e4)) & ~ (e1 = op(e4, e4)) & ~ (e0 = op(e4, e4))) | (~ (e4 = op(e4, e3)) & ~ (e3 = op(e4, e3)) & ~ (e2 = op(e4, e3)) & ~ (e1 = op(e4, e3)) & ~ (e0 = op(e4, e3))) | (~ (e4 = op(e4, e2)) & ~ (e3 = op(e4, e2)) & ~ (e2 = op(e4, e2)) & ~ (e1 = op(e4, e2)) & ~ (e0 = op(e4, e2))) | (~ (e4 = op(e4, e1)) & ~ (e3 = op(e4, e1)) & ~ (e2 = op(e4, e1)) & ~ (e1 = op(e4, e1)) & ~ (e0 = op(e4, e1))) | (~ (e4 = op(e4, e0)) & ~ (e3 = op(e4, e0)) & ~ (e2 = op(e4, e0)) & ~ (e1 = op(e4, e0)) & ~ (e0 = op(e4, e0))) | (~ (e4 = op(e3, e4)) & ~ (e3 = op(e3, e4)) & ~ (e2 = op(e3, e4)) & ~ (e1 = op(e3, e4)) & ~ (e0 = op(e3, e4))) | (~ (e4 = op(e3, e3)) & ~ (e3 = op(e3, e3)) & ~ (e2 = op(e3, e3)) & ~ (e1 = op(e3, e3)) & ~ (e0 = op(e3, e3))) | (~ (e4 = op(e3, e2)) & ~ (e3 = op(e3, e2)) & ~ (e2 = op(e3, e2)) & ~ (e1 = op(e3, e2)) & ~ (e0 = op(e3, e2))) | (~ (e4 = op(e3, e1)) & ~ (e3 = op(e3, e1)) & ~ (e2 = op(e3, e1)) & ~ (e1 = op(e3, e1)) & ~ (e0 = op(e3, e1))) | (~ (e4 = op(e3, e0)) & ~ (e3 = op(e3, e0)) & ~ (e2 = op(e3, e0)) & ~ (e1 = op(e3, e0)) & ~ (e0 = op(e3, e0))) | (~ (e4 = op(e2, e4)) & ~ (e3 = op(e2, e4)) & ~ (e2 = op(e2, e4)) & ~ (e1 = op(e2, e4)) & ~ (e0 = op(e2, e4))) | (~ (e4 = op(e2, e3)) & ~ (e3 = op(e2, e3)) & ~ (e2 = op(e2, e3)) & ~ (e1 = op(e2, e3)) & ~ (e0 = op(e2, e3))) | (~ (e4 = op(e2, e2)) & ~ (e3 = op(e2, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e2, e2)) & ~ (e0 = op(e2, e2))) | (~ (e4 = op(e2, e1)) & ~ (e3 = op(e2, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e2, e1)) & ~ (e0 = op(e2, e1))) | (~ (e4 = op(e2, e0)) & ~ (e3 = op(e2, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e2, e0)) & ~ (e0 = op(e2, e0))) | (~ (e4 = op(e1, e4)) & ~ (e3 = op(e1, e4)) & ~ (e2 = op(e1, e4)) & ~ (e1 = op(e1, e4)) & ~ (e0 = op(e1, e4))) | (~ (e4 = op(e1, e3)) & ~ (e3 = op(e1, e3)) & ~ (e2 = op(e1, e3)) & ~ (e1 = op(e1, e3)) & ~ (e0 = op(e1, e3))) | (~ (e4 = op(e1, e2)) & ~ (e3 = op(e1, e2)) & ~ (e2 = op(e1, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e1, e2))) | (~ (e4 = op(e1, e1)) & ~ (e3 = op(e1, e1)) & ~ (e2 = op(e1, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e1, e1))) | (~ (e4 = op(e1, e0)) & ~ (e3 = op(e1, e0)) & ~ (e2 = op(e1, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e1, e0))) | (~ (e4 = op(e0, e4)) & ~ (e3 = op(e0, e4)) & ~ (e2 = op(e0, e4)) & ~ (e1 = op(e0, e4)) & ~ (e0 = op(e0, e4))) | (~ (e4 = op(e0, e3)) & ~ (e3 = op(e0, e3)) & ~ (e2 = op(e0, e3)) & ~ (e1 = op(e0, e3)) & ~ (e0 = op(e0, e3))) | (~ (e4 = op(e0, e2)) & ~ (e3 = op(e0, e2)) & ~ (e2 = op(e0, e2)) & ~ (e1 = op(e0, e2)) & ~ (e0 = op(e0, e2))) | (~ (e4 = op(e0, e1)) & ~ (e3 = op(e0, e1)) & ~ (e2 = op(e0, e1)) & ~ (e1 = op(e0, e1)) & ~ (e0 = op(e0, e1))) | (~ (e4 = op(e0, e0)) & ~ (e3 = op(e0, e0)) & ~ (e2 = op(e0, e0)) & ~ (e1 = op(e0, e0)) & ~ (e0 = op(e0, e0))) | ~ (e4 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, op(e4, e3))) | ~ (e2 = op(e4, op(e4, e2))) | ~ (e1 = op(e4, op(e4, e1))) | ~ (e0 = op(e4, op(e4, e0))) | ~ (e4 = op(e3, op(e3, e4))) | ~ (e3 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, op(e3, e2))) | ~ (e1 = op(e3, op(e3, e1))) | ~ (e0 = op(e3, op(e3, e0))) | ~ (e4 = op(e2, op(e2, e4))) | ~ (e3 = op(e2, op(e2, e3))) | ~ (e2 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, op(e2, e1))) | ~ (e0 = op(e2, op(e2, e0))) | ~ (e4 = op(e1, op(e1, e4))) | ~ (e3 = op(e1, op(e1, e3))) | ~ (e2 = op(e1, op(e1, e2))) | ~ (e1 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, op(e1, e0))) | ~ (e4 = op(e0, op(e0, e4))) | ~ (e3 = op(e0, op(e0, e3))) | ~ (e2 = op(e0, op(e0, e2))) | ~ (e1 = op(e0, op(e0, e1))) | ~ (e0 = op(e0, op(e0, e0)))), inference(ennf_transformation, [], [f4])).
fof(f4, plain, ~ ((e4 = op(op(e4, e4), op(e4, e4))) & (e4 = op(op(e3, e4), op(e4, e3))) & (e4 = op(op(e2, e4), op(e4, e2))) & (e4 = op(op(e1, e4), op(e4, e1))) & (e4 = op(op(e0, e4), op(e4, e0))) & (e3 = op(op(e4, e3), op(e3, e4))) & (e3 = op(op(e3, e3), op(e3, e3))) & (e3 = op(op(e2, e3), op(e3, e2))) & (e3 = op(op(e1, e3), op(e3, e1))) & (e3 = op(op(e0, e3), op(e3, e0))) & (e2 = op(op(e4, e2), op(e2, e4))) & (e2 = op(op(e3, e2), op(e2, e3))) & (e2 = op(op(e2, e2), op(e2, e2))) & (e2 = op(op(e1, e2), op(e2, e1))) & (e2 = op(op(e0, e2), op(e2, e0))) & (e1 = op(op(e4, e1), op(e1, e4))) & (e1 = op(op(e3, e1), op(e1, e3))) & (e1 = op(op(e2, e1), op(e1, e2))) & (e1 = op(op(e1, e1), op(e1, e1))) & (e1 = op(op(e0, e1), op(e1, e0))) & (e0 = op(op(e4, e0), op(e0, e4))) & (e0 = op(op(e3, e0), op(e0, e3))) & (e0 = op(op(e2, e0), op(e0, e2))) & (e0 = op(op(e1, e0), op(e0, e1))) & (e0 = op(op(e0, e0), op(e0, e0))) & ((e4 = op(e4, e4)) | (e4 = op(e3, e4)) | (e4 = op(e2, e4)) | (e4 = op(e1, e4)) | (e4 = op(e0, e4))) & ((e4 = op(e4, e4)) | (e4 = op(e4, e3)) | (e4 = op(e4, e2)) | (e4 = op(e4, e1)) | (e4 = op(e4, e0))) & ((e3 = op(e4, e4)) | (e3 = op(e3, e4)) | (e3 = op(e2, e4)) | (e3 = op(e1, e4)) | (e3 = op(e0, e4))) & ((e3 = op(e4, e4)) | (e3 = op(e4, e3)) | (e3 = op(e4, e2)) | (e3 = op(e4, e1)) | (e3 = op(e4, e0))) & ((e2 = op(e4, e4)) | (e2 = op(e3, e4)) | (e2 = op(e2, e4)) | (e2 = op(e1, e4)) | (e2 = op(e0, e4))) & ((e2 = op(e4, e4)) | (e2 = op(e4, e3)) | (e2 = op(e4, e2)) | (e2 = op(e4, e1)) | (e2 = op(e4, e0))) & ((e1 = op(e4, e4)) | (e1 = op(e3, e4)) | (e1 = op(e2, e4)) | (e1 = op(e1, e4)) | (e1 = op(e0, e4))) & ((e1 = op(e4, e4)) | (e1 = op(e4, e3)) | (e1 = op(e4, e2)) | (e1 = op(e4, e1)) | (e1 = op(e4, e0))) & ((e0 = op(e4, e4)) | (e0 = op(e3, e4)) | (e0 = op(e2, e4)) | (e0 = op(e1, e4)) | (e0 = op(e0, e4))) & ((e0 = op(e4, e4)) | (e0 = op(e4, e3)) | (e0 = op(e4, e2)) | (e0 = op(e4, e1)) | (e0 = op(e4, e0))) & ((e4 = op(e4, e3)) | (e4 = op(e3, e3)) | (e4 = op(e2, e3)) | (e4 = op(e1, e3)) | (e4 = op(e0, e3))) & ((e4 = op(e3, e4)) | (e4 = op(e3, e3)) | (e4 = op(e3, e2)) | (e4 = op(e3, e1)) | (e4 = op(e3, e0))) & ((e3 = op(e4, e3)) | (e3 = op(e3, e3)) | (e3 = op(e2, e3)) | (e3 = op(e1, e3)) | (e3 = op(e0, e3))) & ((e3 = op(e3, e4)) | (e3 = op(e3, e3)) | (e3 = op(e3, e2)) | (e3 = op(e3, e1)) | (e3 = op(e3, e0))) & ((e2 = op(e4, e3)) | (e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(e0, e3))) & ((e2 = op(e3, e4)) | (e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, e0))) & ((e1 = op(e4, e3)) | (e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(e0, e3))) & ((e1 = op(e3, e4)) | (e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, e0))) & ((e0 = op(e4, e3)) | (e0 = op(e3, e3)) | (e0 = op(e2, e3)) | (e0 = op(e1, e3)) | (e0 = op(e0, e3))) & ((e0 = op(e3, e4)) | (e0 = op(e3, e3)) | (e0 = op(e3, e2)) | (e0 = op(e3, e1)) | (e0 = op(e3, e0))) & ((e4 = op(e4, e2)) | (e4 = op(e3, e2)) | (e4 = op(e2, e2)) | (e4 = op(e1, e2)) | (e4 = op(e0, e2))) & ((e4 = op(e2, e4)) | (e4 = op(e2, e3)) | (e4 = op(e2, e2)) | (e4 = op(e2, e1)) | (e4 = op(e2, e0))) & ((e3 = op(e4, e2)) | (e3 = op(e3, e2)) | (e3 = op(e2, e2)) | (e3 = op(e1, e2)) | (e3 = op(e0, e2))) & ((e3 = op(e2, e4)) | (e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, e0))) & ((e2 = op(e4, e2)) | (e2 = op(e3, e2)) | (e2 = op(e2, e2)) | (e2 = op(e1, e2)) | (e2 = op(e0, e2))) & ((e2 = op(e2, e4)) | (e2 = op(e2, e3)) | (e2 = op(e2, e2)) | (e2 = op(e2, e1)) | (e2 = op(e2, e0))) & ((e1 = op(e4, e2)) | (e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(e0, e2))) & ((e1 = op(e2, e4)) | (e1 = op(e2, e3)) | (e1 = op(e2, e2)) | (e1 = op(e2, e1)) | (e1 = op(e2, e0))) & ((e0 = op(e4, e2)) | (e0 = op(e3, e2)) | (e0 = op(e2, e2)) | (e0 = op(e1, e2)) | (e0 = op(e0, e2))) & ((e0 = op(e2, e4)) | (e0 = op(e2, e3)) | (e0 = op(e2, e2)) | (e0 = op(e2, e1)) | (e0 = op(e2, e0))) & ((e4 = op(e4, e1)) | (e4 = op(e3, e1)) | (e4 = op(e2, e1)) | (e4 = op(e1, e1)) | (e4 = op(e0, e1))) & ((e4 = op(e1, e4)) | (e4 = op(e1, e3)) | (e4 = op(e1, e2)) | (e4 = op(e1, e1)) | (e4 = op(e1, e0))) & ((e3 = op(e4, e1)) | (e3 = op(e3, e1)) | (e3 = op(e2, e1)) | (e3 = op(e1, e1)) | (e3 = op(e0, e1))) & ((e3 = op(e1, e4)) | (e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, e0))) & ((e2 = op(e4, e1)) | (e2 = op(e3, e1)) | (e2 = op(e2, e1)) | (e2 = op(e1, e1)) | (e2 = op(e0, e1))) & ((e2 = op(e1, e4)) | (e2 = op(e1, e3)) | (e2 = op(e1, e2)) | (e2 = op(e1, e1)) | (e2 = op(e1, e0))) & ((e1 = op(e4, e1)) | (e1 = op(e3, e1)) | (e1 = op(e2, e1)) | (e1 = op(e1, e1)) | (e1 = op(e0, e1))) & ((e1 = op(e1, e4)) | (e1 = op(e1, e3)) | (e1 = op(e1, e2)) | (e1 = op(e1, e1)) | (e1 = op(e1, e0))) & ((e0 = op(e4, e1)) | (e0 = op(e3, e1)) | (e0 = op(e2, e1)) | (e0 = op(e1, e1)) | (e0 = op(e0, e1))) & ((e0 = op(e1, e4)) | (e0 = op(e1, e3)) | (e0 = op(e1, e2)) | (e0 = op(e1, e1)) | (e0 = op(e1, e0))) & ((e4 = op(e4, e0)) | (e4 = op(e3, e0)) | (e4 = op(e2, e0)) | (e4 = op(e1, e0)) | (e4 = op(e0, e0))) & ((e4 = op(e0, e4)) | (e4 = op(e0, e3)) | (e4 = op(e0, e2)) | (e4 = op(e0, e1)) | (e4 = op(e0, e0))) & ((e3 = op(e4, e0)) | (e3 = op(e3, e0)) | (e3 = op(e2, e0)) | (e3 = op(e1, e0)) | (e3 = op(e0, e0))) & ((e3 = op(e0, e4)) | (e3 = op(e0, e3)) | (e3 = op(e0, e2)) | (e3 = op(e0, e1)) | (e3 = op(e0, e0))) & ((e2 = op(e4, e0)) | (e2 = op(e3, e0)) | (e2 = op(e2, e0)) | (e2 = op(e1, e0)) | (e2 = op(e0, e0))) & ((e2 = op(e0, e4)) | (e2 = op(e0, e3)) | (e2 = op(e0, e2)) | (e2 = op(e0, e1)) | (e2 = op(e0, e0))) & ((e1 = op(e4, e0)) | (e1 = op(e3, e0)) | (e1 = op(e2, e0)) | (e1 = op(e1, e0)) | (e1 = op(e0, e0))) & ((e1 = op(e0, e4)) | (e1 = op(e0, e3)) | (e1 = op(e0, e2)) | (e1 = op(e0, e1)) | (e1 = op(e0, e0))) & ((e0 = op(e4, e0)) | (e0 = op(e3, e0)) | (e0 = op(e2, e0)) | (e0 = op(e1, e0)) | (e0 = op(e0, e0))) & ((e0 = op(e0, e4)) | (e0 = op(e0, e3)) | (e0 = op(e0, e2)) | (e0 = op(e0, e1)) | (e0 = op(e0, e0))) & ((e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | (e1 = op(e4, e4)) | (e0 = op(e4, e4))) & ((e4 = op(e4, e3)) | (e3 = op(e4, e3)) | (e2 = op(e4, e3)) | (e1 = op(e4, e3)) | (e0 = op(e4, e3))) & ((e4 = op(e4, e2)) | (e3 = op(e4, e2)) | (e2 = op(e4, e2)) | (e1 = op(e4, e2)) | (e0 = op(e4, e2))) & ((e4 = op(e4, e1)) | (e3 = op(e4, e1)) | (e2 = op(e4, e1)) | (e1 = op(e4, e1)) | (e0 = op(e4, e1))) & ((e4 = op(e4, e0)) | (e3 = op(e4, e0)) | (e2 = op(e4, e0)) | (e1 = op(e4, e0)) | (e0 = op(e4, e0))) & ((e4 = op(e3, e4)) | (e3 = op(e3, e4)) | (e2 = op(e3, e4)) | (e1 = op(e3, e4)) | (e0 = op(e3, e4))) & ((e4 = op(e3, e3)) | (e3 = op(e3, e3)) | (e2 = op(e3, e3)) | (e1 = op(e3, e3)) | (e0 = op(e3, e3))) & ((e4 = op(e3, e2)) | (e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (e0 = op(e3, e2))) & ((e4 = op(e3, e1)) | (e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (e0 = op(e3, e1))) & ((e4 = op(e3, e0)) | (e3 = op(e3, e0)) | (e2 = op(e3, e0)) | (e1 = op(e3, e0)) | (e0 = op(e3, e0))) & ((e4 = op(e2, e4)) | (e3 = op(e2, e4)) | (e2 = op(e2, e4)) | (e1 = op(e2, e4)) | (e0 = op(e2, e4))) & ((e4 = op(e2, e3)) | (e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))) & ((e4 = op(e2, e2)) | (e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (e0 = op(e2, e2))) & ((e4 = op(e2, e1)) | (e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))) & ((e4 = op(e2, e0)) | (e3 = op(e2, e0)) | (e2 = op(e2, e0)) | (e1 = op(e2, e0)) | (e0 = op(e2, e0))) & ((e4 = op(e1, e4)) | (e3 = op(e1, e4)) | (e2 = op(e1, e4)) | (e1 = op(e1, e4)) | (e0 = op(e1, e4))) & ((e4 = op(e1, e3)) | (e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (e0 = op(e1, e3))) & ((e4 = op(e1, e2)) | (e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (e0 = op(e1, e2))) & ((e4 = op(e1, e1)) | (e3 = op(e1, e1)) | (e2 = op(e1, e1)) | (e1 = op(e1, e1)) | (e0 = op(e1, e1))) & ((e4 = op(e1, e0)) | (e3 = op(e1, e0)) | (e2 = op(e1, e0)) | (e1 = op(e1, e0)) | (e0 = op(e1, e0))) & ((e4 = op(e0, e4)) | (e3 = op(e0, e4)) | (e2 = op(e0, e4)) | (e1 = op(e0, e4)) | (e0 = op(e0, e4))) & ((e4 = op(e0, e3)) | (e3 = op(e0, e3)) | (e2 = op(e0, e3)) | (e1 = op(e0, e3)) | (e0 = op(e0, e3))) & ((e4 = op(e0, e2)) | (e3 = op(e0, e2)) | (e2 = op(e0, e2)) | (e1 = op(e0, e2)) | (e0 = op(e0, e2))) & ((e4 = op(e0, e1)) | (e3 = op(e0, e1)) | (e2 = op(e0, e1)) | (e1 = op(e0, e1)) | (e0 = op(e0, e1))) & ((e4 = op(e0, e0)) | (e3 = op(e0, e0)) | (e2 = op(e0, e0)) | (e1 = op(e0, e0)) | (e0 = op(e0, e0))) & (e4 = op(e4, op(e4, e4))) & (e3 = op(e4, op(e4, e3))) & (e2 = op(e4, op(e4, e2))) & (e1 = op(e4, op(e4, e1))) & (e0 = op(e4, op(e4, e0))) & (e4 = op(e3, op(e3, e4))) & (e3 = op(e3, op(e3, e3))) & (e2 = op(e3, op(e3, e2))) & (e1 = op(e3, op(e3, e1))) & (e0 = op(e3, op(e3, e0))) & (e4 = op(e2, op(e2, e4))) & (e3 = op(e2, op(e2, e3))) & (e2 = op(e2, op(e2, e2))) & (e1 = op(e2, op(e2, e1))) & (e0 = op(e2, op(e2, e0))) & (e4 = op(e1, op(e1, e4))) & (e3 = op(e1, op(e1, e3))) & (e2 = op(e1, op(e1, e2))) & (e1 = op(e1, op(e1, e1))) & (e0 = op(e1, op(e1, e0))) & (e4 = op(e0, op(e0, e4))) & (e3 = op(e0, op(e0, e3))) & (e2 = op(e0, op(e0, e2))) & (e1 = op(e0, op(e0, e1))) & (e0 = op(e0, op(e0, e0)))), inference(negated_conjecture, [], [f3])).
fof(f3, plain, ~ ((e4 = op(op(e4, e4), op(e4, e4))) & (e4 = op(op(e3, e4), op(e4, e3))) & (e4 = op(op(e2, e4), op(e4, e2))) & (e4 = op(op(e1, e4), op(e4, e1))) & (e4 = op(op(e0, e4), op(e4, e0))) & (e3 = op(op(e4, e3), op(e3, e4))) & (e3 = op(op(e3, e3), op(e3, e3))) & (e3 = op(op(e2, e3), op(e3, e2))) & (e3 = op(op(e1, e3), op(e3, e1))) & (e3 = op(op(e0, e3), op(e3, e0))) & (e2 = op(op(e4, e2), op(e2, e4))) & (e2 = op(op(e3, e2), op(e2, e3))) & (e2 = op(op(e2, e2), op(e2, e2))) & (e2 = op(op(e1, e2), op(e2, e1))) & (e2 = op(op(e0, e2), op(e2, e0))) & (e1 = op(op(e4, e1), op(e1, e4))) & (e1 = op(op(e3, e1), op(e1, e3))) & (e1 = op(op(e2, e1), op(e1, e2))) & (e1 = op(op(e1, e1), op(e1, e1))) & (e1 = op(op(e0, e1), op(e1, e0))) & (e0 = op(op(e4, e0), op(e0, e4))) & (e0 = op(op(e3, e0), op(e0, e3))) & (e0 = op(op(e2, e0), op(e0, e2))) & (e0 = op(op(e1, e0), op(e0, e1))) & (e0 = op(op(e0, e0), op(e0, e0))) & ((e4 = op(e4, e4)) | (e4 = op(e3, e4)) | (e4 = op(e2, e4)) | (e4 = op(e1, e4)) | (e4 = op(e0, e4))) & ((e4 = op(e4, e4)) | (e4 = op(e4, e3)) | (e4 = op(e4, e2)) | (e4 = op(e4, e1)) | (e4 = op(e4, e0))) & ((e3 = op(e4, e4)) | (e3 = op(e3, e4)) | (e3 = op(e2, e4)) | (e3 = op(e1, e4)) | (e3 = op(e0, e4))) & ((e3 = op(e4, e4)) | (e3 = op(e4, e3)) | (e3 = op(e4, e2)) | (e3 = op(e4, e1)) | (e3 = op(e4, e0))) & ((e2 = op(e4, e4)) | (e2 = op(e3, e4)) | (e2 = op(e2, e4)) | (e2 = op(e1, e4)) | (e2 = op(e0, e4))) & ((e2 = op(e4, e4)) | (e2 = op(e4, e3)) | (e2 = op(e4, e2)) | (e2 = op(e4, e1)) | (e2 = op(e4, e0))) & ((e1 = op(e4, e4)) | (e1 = op(e3, e4)) | (e1 = op(e2, e4)) | (e1 = op(e1, e4)) | (e1 = op(e0, e4))) & ((e1 = op(e4, e4)) | (e1 = op(e4, e3)) | (e1 = op(e4, e2)) | (e1 = op(e4, e1)) | (e1 = op(e4, e0))) & ((e0 = op(e4, e4)) | (e0 = op(e3, e4)) | (e0 = op(e2, e4)) | (e0 = op(e1, e4)) | (e0 = op(e0, e4))) & ((e0 = op(e4, e4)) | (e0 = op(e4, e3)) | (e0 = op(e4, e2)) | (e0 = op(e4, e1)) | (e0 = op(e4, e0))) & ((e4 = op(e4, e3)) | (e4 = op(e3, e3)) | (e4 = op(e2, e3)) | (e4 = op(e1, e3)) | (e4 = op(e0, e3))) & ((e4 = op(e3, e4)) | (e4 = op(e3, e3)) | (e4 = op(e3, e2)) | (e4 = op(e3, e1)) | (e4 = op(e3, e0))) & ((e3 = op(e4, e3)) | (e3 = op(e3, e3)) | (e3 = op(e2, e3)) | (e3 = op(e1, e3)) | (e3 = op(e0, e3))) & ((e3 = op(e3, e4)) | (e3 = op(e3, e3)) | (e3 = op(e3, e2)) | (e3 = op(e3, e1)) | (e3 = op(e3, e0))) & ((e2 = op(e4, e3)) | (e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(e0, e3))) & ((e2 = op(e3, e4)) | (e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, e0))) & ((e1 = op(e4, e3)) | (e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(e0, e3))) & ((e1 = op(e3, e4)) | (e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, e0))) & ((e0 = op(e4, e3)) | (e0 = op(e3, e3)) | (e0 = op(e2, e3)) | (e0 = op(e1, e3)) | (e0 = op(e0, e3))) & ((e0 = op(e3, e4)) | (e0 = op(e3, e3)) | (e0 = op(e3, e2)) | (e0 = op(e3, e1)) | (e0 = op(e3, e0))) & ((e4 = op(e4, e2)) | (e4 = op(e3, e2)) | (e4 = op(e2, e2)) | (e4 = op(e1, e2)) | (e4 = op(e0, e2))) & ((e4 = op(e2, e4)) | (e4 = op(e2, e3)) | (e4 = op(e2, e2)) | (e4 = op(e2, e1)) | (e4 = op(e2, e0))) & ((e3 = op(e4, e2)) | (e3 = op(e3, e2)) | (e3 = op(e2, e2)) | (e3 = op(e1, e2)) | (e3 = op(e0, e2))) & ((e3 = op(e2, e4)) | (e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, e0))) & ((e2 = op(e4, e2)) | (e2 = op(e3, e2)) | (e2 = op(e2, e2)) | (e2 = op(e1, e2)) | (e2 = op(e0, e2))) & ((e2 = op(e2, e4)) | (e2 = op(e2, e3)) | (e2 = op(e2, e2)) | (e2 = op(e2, e1)) | (e2 = op(e2, e0))) & ((e1 = op(e4, e2)) | (e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(e0, e2))) & ((e1 = op(e2, e4)) | (e1 = op(e2, e3)) | (e1 = op(e2, e2)) | (e1 = op(e2, e1)) | (e1 = op(e2, e0))) & ((e0 = op(e4, e2)) | (e0 = op(e3, e2)) | (e0 = op(e2, e2)) | (e0 = op(e1, e2)) | (e0 = op(e0, e2))) & ((e0 = op(e2, e4)) | (e0 = op(e2, e3)) | (e0 = op(e2, e2)) | (e0 = op(e2, e1)) | (e0 = op(e2, e0))) & ((e4 = op(e4, e1)) | (e4 = op(e3, e1)) | (e4 = op(e2, e1)) | (e4 = op(e1, e1)) | (e4 = op(e0, e1))) & ((e4 = op(e1, e4)) | (e4 = op(e1, e3)) | (e4 = op(e1, e2)) | (e4 = op(e1, e1)) | (e4 = op(e1, e0))) & ((e3 = op(e4, e1)) | (e3 = op(e3, e1)) | (e3 = op(e2, e1)) | (e3 = op(e1, e1)) | (e3 = op(e0, e1))) & ((e3 = op(e1, e4)) | (e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, e0))) & ((e2 = op(e4, e1)) | (e2 = op(e3, e1)) | (e2 = op(e2, e1)) | (e2 = op(e1, e1)) | (e2 = op(e0, e1))) & ((e2 = op(e1, e4)) | (e2 = op(e1, e3)) | (e2 = op(e1, e2)) | (e2 = op(e1, e1)) | (e2 = op(e1, e0))) & ((e1 = op(e4, e1)) | (e1 = op(e3, e1)) | (e1 = op(e2, e1)) | (e1 = op(e1, e1)) | (e1 = op(e0, e1))) & ((e1 = op(e1, e4)) | (e1 = op(e1, e3)) | (e1 = op(e1, e2)) | (e1 = op(e1, e1)) | (e1 = op(e1, e0))) & ((e0 = op(e4, e1)) | (e0 = op(e3, e1)) | (e0 = op(e2, e1)) | (e0 = op(e1, e1)) | (e0 = op(e0, e1))) & ((e0 = op(e1, e4)) | (e0 = op(e1, e3)) | (e0 = op(e1, e2)) | (e0 = op(e1, e1)) | (e0 = op(e1, e0))) & ((e4 = op(e4, e0)) | (e4 = op(e3, e0)) | (e4 = op(e2, e0)) | (e4 = op(e1, e0)) | (e4 = op(e0, e0))) & ((e4 = op(e0, e4)) | (e4 = op(e0, e3)) | (e4 = op(e0, e2)) | (e4 = op(e0, e1)) | (e4 = op(e0, e0))) & ((e3 = op(e4, e0)) | (e3 = op(e3, e0)) | (e3 = op(e2, e0)) | (e3 = op(e1, e0)) | (e3 = op(e0, e0))) & ((e3 = op(e0, e4)) | (e3 = op(e0, e3)) | (e3 = op(e0, e2)) | (e3 = op(e0, e1)) | (e3 = op(e0, e0))) & ((e2 = op(e4, e0)) | (e2 = op(e3, e0)) | (e2 = op(e2, e0)) | (e2 = op(e1, e0)) | (e2 = op(e0, e0))) & ((e2 = op(e0, e4)) | (e2 = op(e0, e3)) | (e2 = op(e0, e2)) | (e2 = op(e0, e1)) | (e2 = op(e0, e0))) & ((e1 = op(e4, e0)) | (e1 = op(e3, e0)) | (e1 = op(e2, e0)) | (e1 = op(e1, e0)) | (e1 = op(e0, e0))) & ((e1 = op(e0, e4)) | (e1 = op(e0, e3)) | (e1 = op(e0, e2)) | (e1 = op(e0, e1)) | (e1 = op(e0, e0))) & ((e0 = op(e4, e0)) | (e0 = op(e3, e0)) | (e0 = op(e2, e0)) | (e0 = op(e1, e0)) | (e0 = op(e0, e0))) & ((e0 = op(e0, e4)) | (e0 = op(e0, e3)) | (e0 = op(e0, e2)) | (e0 = op(e0, e1)) | (e0 = op(e0, e0))) & ((e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | (e1 = op(e4, e4)) | (e0 = op(e4, e4))) & ((e4 = op(e4, e3)) | (e3 = op(e4, e3)) | (e2 = op(e4, e3)) | (e1 = op(e4, e3)) | (e0 = op(e4, e3))) & ((e4 = op(e4, e2)) | (e3 = op(e4, e2)) | (e2 = op(e4, e2)) | (e1 = op(e4, e2)) | (e0 = op(e4, e2))) & ((e4 = op(e4, e1)) | (e3 = op(e4, e1)) | (e2 = op(e4, e1)) | (e1 = op(e4, e1)) | (e0 = op(e4, e1))) & ((e4 = op(e4, e0)) | (e3 = op(e4, e0)) | (e2 = op(e4, e0)) | (e1 = op(e4, e0)) | (e0 = op(e4, e0))) & ((e4 = op(e3, e4)) | (e3 = op(e3, e4)) | (e2 = op(e3, e4)) | (e1 = op(e3, e4)) | (e0 = op(e3, e4))) & ((e4 = op(e3, e3)) | (e3 = op(e3, e3)) | (e2 = op(e3, e3)) | (e1 = op(e3, e3)) | (e0 = op(e3, e3))) & ((e4 = op(e3, e2)) | (e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (e0 = op(e3, e2))) & ((e4 = op(e3, e1)) | (e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (e0 = op(e3, e1))) & ((e4 = op(e3, e0)) | (e3 = op(e3, e0)) | (e2 = op(e3, e0)) | (e1 = op(e3, e0)) | (e0 = op(e3, e0))) & ((e4 = op(e2, e4)) | (e3 = op(e2, e4)) | (e2 = op(e2, e4)) | (e1 = op(e2, e4)) | (e0 = op(e2, e4))) & ((e4 = op(e2, e3)) | (e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))) & ((e4 = op(e2, e2)) | (e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (e0 = op(e2, e2))) & ((e4 = op(e2, e1)) | (e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))) & ((e4 = op(e2, e0)) | (e3 = op(e2, e0)) | (e2 = op(e2, e0)) | (e1 = op(e2, e0)) | (e0 = op(e2, e0))) & ((e4 = op(e1, e4)) | (e3 = op(e1, e4)) | (e2 = op(e1, e4)) | (e1 = op(e1, e4)) | (e0 = op(e1, e4))) & ((e4 = op(e1, e3)) | (e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (e0 = op(e1, e3))) & ((e4 = op(e1, e2)) | (e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (e0 = op(e1, e2))) & ((e4 = op(e1, e1)) | (e3 = op(e1, e1)) | (e2 = op(e1, e1)) | (e1 = op(e1, e1)) | (e0 = op(e1, e1))) & ((e4 = op(e1, e0)) | (e3 = op(e1, e0)) | (e2 = op(e1, e0)) | (e1 = op(e1, e0)) | (e0 = op(e1, e0))) & ((e4 = op(e0, e4)) | (e3 = op(e0, e4)) | (e2 = op(e0, e4)) | (e1 = op(e0, e4)) | (e0 = op(e0, e4))) & ((e4 = op(e0, e3)) | (e3 = op(e0, e3)) | (e2 = op(e0, e3)) | (e1 = op(e0, e3)) | (e0 = op(e0, e3))) & ((e4 = op(e0, e2)) | (e3 = op(e0, e2)) | (e2 = op(e0, e2)) | (e1 = op(e0, e2)) | (e0 = op(e0, e2))) & ((e4 = op(e0, e1)) | (e3 = op(e0, e1)) | (e2 = op(e0, e1)) | (e1 = op(e0, e1)) | (e0 = op(e0, e1))) & ((e4 = op(e0, e0)) | (e3 = op(e0, e0)) | (e2 = op(e0, e0)) | (e1 = op(e0, e0)) | (e0 = op(e0, e0))) & (e4 = op(e4, op(e4, e4))) & (e3 = op(e4, op(e4, e3))) & (e2 = op(e4, op(e4, e2))) & (e1 = op(e4, op(e4, e1))) & (e0 = op(e4, op(e4, e0))) & (e4 = op(e3, op(e3, e4))) & (e3 = op(e3, op(e3, e3))) & (e2 = op(e3, op(e3, e2))) & (e1 = op(e3, op(e3, e1))) & (e0 = op(e3, op(e3, e0))) & (e4 = op(e2, op(e2, e4))) & (e3 = op(e2, op(e2, e3))) & (e2 = op(e2, op(e2, e2))) & (e1 = op(e2, op(e2, e1))) & (e0 = op(e2, op(e2, e0))) & (e4 = op(e1, op(e1, e4))) & (e3 = op(e1, op(e1, e3))) & (e2 = op(e1, op(e1, e2))) & (e1 = op(e1, op(e1, e1))) & (e0 = op(e1, op(e1, e0))) & (e4 = op(e0, op(e0, e4))) & (e3 = op(e0, op(e0, e3))) & (e2 = op(e0, op(e0, e2))) & (e1 = op(e0, op(e0, e1))) & (e0 = op(e0, op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG188+1.p', co1)).