fof(f1386, plain, $false, inference(subsumption_resolution, [], [f1385, f191])).
fof(f191, plain, (e0 = op(e4, e4)), inference(cnf_transformation, [], [f2])).
fof(f2, plain, ((e0 = op(e4, e4)) & (e4 = op(e4, e3)) & (e1 = op(e4, e2)) & (e3 = op(e4, e1)) & (e2 = op(e4, e0)) & (e1 = op(e3, e4)) & (e2 = op(e3, e3)) & (e4 = op(e3, e2)) & (e0 = op(e3, e1)) & (e3 = op(e3, e0)) & (e2 = op(e2, e4)) & (e0 = op(e2, e3)) & (e3 = op(e2, e2)) & (e4 = op(e2, e1)) & (e1 = op(e2, e0)) & (e4 = op(e1, e4)) & (e3 = op(e1, e3)) & (e2 = op(e1, e2)) & (e1 = op(e1, e1)) & (e0 = op(e1, e0)) & (e3 = op(e0, e4)) & (e1 = op(e0, e3)) & (e0 = op(e0, e2)) & (e2 = op(e0, e1)) & (e4 = op(e0, e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG187+1.p', ax2)).
fof(f1385, plain, ~ (e0 = op(e4, e4)), inference(forward_demodulation, [], [f1384, f167])).
fof(f167, plain, (e4 = op(e0, e0)), inference(cnf_transformation, [], [f2])).
fof(f1384, plain, ~ (e0 = op(op(e0, e0), op(e0, e0))), inference(subsumption_resolution, [], [f1383, f169])).
fof(f169, plain, (e0 = op(e0, e2)), inference(cnf_transformation, [], [f2])).
fof(f1383, plain, (~ (e0 = op(e0, e2)) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(forward_demodulation, [], [f1382, f172])).
fof(f172, plain, (e0 = op(e1, e0)), inference(cnf_transformation, [], [f2])).
fof(f1382, plain, (~ (e0 = op(op(e1, e0), e2)) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(forward_demodulation, [], [f1381, f168])).
fof(f168, plain, (e2 = op(e0, e1)), inference(cnf_transformation, [], [f2])).
fof(f1381, plain, (~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(subsumption_resolution, [], [f1380, f172])).
fof(f1380, plain, (~ (e0 = op(e1, e0)) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(forward_demodulation, [], [f1379, f177])).
fof(f177, plain, (e1 = op(e2, e0)), inference(cnf_transformation, [], [f2])).
fof(f1379, plain, (~ (e0 = op(op(e2, e0), e0)) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(forward_demodulation, [], [f1378, f169])).
fof(f1378, plain, (~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(subsumption_resolution, [], [f1377, f183])).
fof(f183, plain, (e0 = op(e3, e1)), inference(cnf_transformation, [], [f2])).
fof(f1377, plain, (~ (e0 = op(e3, e1)) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(forward_demodulation, [], [f1376, f182])).
fof(f182, plain, (e3 = op(e3, e0)), inference(cnf_transformation, [], [f2])).
fof(f1376, plain, (~ (e0 = op(op(e3, e0), e1)) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(forward_demodulation, [], [f1375, f170])).
fof(f170, plain, (e1 = op(e0, e3)), inference(cnf_transformation, [], [f2])).
fof(f1375, plain, (~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(subsumption_resolution, [], [f1374, f180])).
fof(f180, plain, (e0 = op(e2, e3)), inference(cnf_transformation, [], [f2])).
fof(f1374, plain, (~ (e0 = op(e2, e3)) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(forward_demodulation, [], [f1373, f187])).
fof(f187, plain, (e2 = op(e4, e0)), inference(cnf_transformation, [], [f2])).
fof(f1373, plain, (~ (e0 = op(op(e4, e0), e3)) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(forward_demodulation, [], [f1372, f171])).
fof(f171, plain, (e3 = op(e0, e4)), inference(cnf_transformation, [], [f2])).
fof(f1372, plain, (~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(subsumption_resolution, [], [f1371, f177])).
fof(f1371, plain, (~ (e1 = op(e2, e0)) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(forward_demodulation, [], [f1370, f168])).
fof(f1370, plain, (~ (e1 = op(op(e0, e1), e0)) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(forward_demodulation, [], [f1369, f172])).
fof(f1369, plain, (~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(subsumption_resolution, [], [f1368, f173])).
fof(f173, plain, (e1 = op(e1, e1)), inference(cnf_transformation, [], [f2])).
fof(f1368, plain, (~ (e1 = op(e1, e1)) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(forward_demodulation, [], [f1367, f173])).
fof(f1367, plain, (~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(subsumption_resolution, [], [f1366, f189])).
fof(f189, plain, (e1 = op(e4, e2)), inference(cnf_transformation, [], [f2])).
fof(f1366, plain, (~ (e1 = op(e4, e2)) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(forward_demodulation, [], [f1365, f178])).
fof(f178, plain, (e4 = op(e2, e1)), inference(cnf_transformation, [], [f2])).
fof(f1365, plain, (~ (e1 = op(op(e2, e1), e2)) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(forward_demodulation, [], [f1364, f174])).
fof(f174, plain, (e2 = op(e1, e2)), inference(cnf_transformation, [], [f2])).
fof(f1364, plain, (~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(subsumption_resolution, [], [f1363, f170])).
fof(f1363, plain, (~ (e1 = op(e0, e3)) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(forward_demodulation, [], [f1362, f183])).
fof(f1362, plain, (~ (e1 = op(op(e3, e1), e3)) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(forward_demodulation, [], [f1361, f175])).
fof(f175, plain, (e3 = op(e1, e3)), inference(cnf_transformation, [], [f2])).
fof(f1361, plain, (~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(subsumption_resolution, [], [f1360, f186])).
fof(f186, plain, (e1 = op(e3, e4)), inference(cnf_transformation, [], [f2])).
fof(f1360, plain, (~ (e1 = op(e3, e4)) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(forward_demodulation, [], [f1359, f188])).
fof(f188, plain, (e3 = op(e4, e1)), inference(cnf_transformation, [], [f2])).
fof(f1359, plain, (~ (e1 = op(op(e4, e1), e4)) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(forward_demodulation, [], [f1358, f176])).
fof(f176, plain, (e4 = op(e1, e4)), inference(cnf_transformation, [], [f2])).
fof(f1358, plain, (~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(subsumption_resolution, [], [f1357, f168])).
fof(f1357, plain, (~ (e2 = op(e0, e1)) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(forward_demodulation, [], [f1356, f169])).
fof(f1356, plain, (~ (e2 = op(op(e0, e2), e1)) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(forward_demodulation, [], [f1355, f177])).
fof(f1355, plain, (~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(subsumption_resolution, [], [f1354, f181])).
fof(f181, plain, (e2 = op(e2, e4)), inference(cnf_transformation, [], [f2])).
fof(f1354, plain, (~ (e2 = op(e2, e4)) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(forward_demodulation, [], [f1353, f174])).
fof(f1353, plain, (~ (e2 = op(op(e1, e2), e4)) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(forward_demodulation, [], [f1352, f178])).
fof(f1352, plain, (~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(subsumption_resolution, [], [f1351, f185])).
fof(f185, plain, (e2 = op(e3, e3)), inference(cnf_transformation, [], [f2])).
fof(f1351, plain, (~ (e2 = op(e3, e3)) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(forward_demodulation, [], [f1350, f179])).
fof(f179, plain, (e3 = op(e2, e2)), inference(cnf_transformation, [], [f2])).
fof(f1350, plain, (~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(subsumption_resolution, [], [f1349, f187])).
fof(f1349, plain, (~ (e2 = op(e4, e0)) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(forward_demodulation, [], [f1348, f184])).
fof(f184, plain, (e4 = op(e3, e2)), inference(cnf_transformation, [], [f2])).
fof(f1348, plain, (~ (e2 = op(op(e3, e2), e0)) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(forward_demodulation, [], [f1347, f180])).
fof(f1347, plain, (~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(subsumption_resolution, [], [f1346, f174])).
fof(f1346, plain, (~ (e2 = op(e1, e2)) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(forward_demodulation, [], [f1345, f189])).
fof(f1345, plain, (~ (e2 = op(op(e4, e2), e2)) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(forward_demodulation, [], [f1344, f181])).
fof(f1344, plain, (~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(subsumption_resolution, [], [f1343, f175])).
fof(f1343, plain, (~ (e3 = op(e1, e3)) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(forward_demodulation, [], [f1342, f170])).
fof(f1342, plain, (~ (e3 = op(op(e0, e3), e3)) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(forward_demodulation, [], [f1341, f182])).
fof(f1341, plain, (~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(subsumption_resolution, [], [f1340, f182])).
fof(f1340, plain, (~ (e3 = op(e3, e0)) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(forward_demodulation, [], [f1339, f175])).
fof(f1339, plain, (~ (e3 = op(op(e1, e3), e0)) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(forward_demodulation, [], [f1338, f183])).
fof(f1338, plain, (~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(subsumption_resolution, [], [f1337, f171])).
fof(f1337, plain, (~ (e3 = op(e0, e4)) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(forward_demodulation, [], [f1336, f180])).
fof(f1336, plain, (~ (e3 = op(op(e2, e3), e4)) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(forward_demodulation, [], [f1335, f184])).
fof(f1335, plain, (~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(subsumption_resolution, [], [f1334, f179])).
fof(f1334, plain, (~ (e3 = op(e2, e2)) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(forward_demodulation, [], [f1333, f185])).
fof(f1333, plain, (~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(subsumption_resolution, [], [f1332, f188])).
fof(f1332, plain, (~ (e3 = op(e4, e1)) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(forward_demodulation, [], [f1331, f190])).
fof(f190, plain, (e4 = op(e4, e3)), inference(cnf_transformation, [], [f2])).
fof(f1331, plain, (~ (e3 = op(op(e4, e3), e1)) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(forward_demodulation, [], [f1330, f186])).
fof(f1330, plain, (~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(subsumption_resolution, [], [f1329, f184])).
fof(f1329, plain, (~ (e4 = op(e3, e2)) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(forward_demodulation, [], [f1328, f171])).
fof(f1328, plain, (~ (e4 = op(op(e0, e4), e2)) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(forward_demodulation, [], [f1327, f187])).
fof(f1327, plain, (~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(subsumption_resolution, [], [f1326, f190])).
fof(f1326, plain, (~ (e4 = op(e4, e3)) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(forward_demodulation, [], [f1325, f176])).
fof(f1325, plain, (~ (e4 = op(op(e1, e4), e3)) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(forward_demodulation, [], [f1324, f188])).
fof(f1324, plain, (~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(subsumption_resolution, [], [f1323, f178])).
fof(f1323, plain, (~ (e4 = op(e2, e1)) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(forward_demodulation, [], [f1322, f181])).
fof(f1322, plain, (~ (e4 = op(op(e2, e4), e1)) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(forward_demodulation, [], [f1321, f189])).
fof(f1321, plain, (~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(subsumption_resolution, [], [f1320, f176])).
fof(f1320, plain, (~ (e4 = op(e1, e4)) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(forward_demodulation, [], [f1319, f186])).
fof(f1319, plain, (~ (e4 = op(op(e3, e4), e4)) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(forward_demodulation, [], [f1318, f190])).
fof(f1318, plain, (~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(subsumption_resolution, [], [f1317, f167])).
fof(f1317, plain, (~ (e4 = op(e0, e0)) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(forward_demodulation, [], [f1316, f191])).
fof(f1316, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0)))), inference(subsumption_resolution, [], [f1315, f802])).
fof(f802, plain, ~ sP0, inference(subsumption_resolution, [], [f801, f158])).
fof(f158, plain, ~ (e0 = e2), inference(cnf_transformation, [], [f1])).
fof(f1, plain, (~ (e3 = e4) & ~ (e2 = e4) & ~ (e2 = e3) & ~ (e1 = e4) & ~ (e1 = e3) & ~ (e1 = e2) & ~ (e0 = e4) & ~ (e0 = e3) & ~ (e0 = e2) & ~ (e0 = e1)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG187+1.p', ax1)).
fof(f801, plain, ((e0 = e2) | ~ sP0), inference(forward_demodulation, [], [f800, f187])).
fof(f800, plain, ((e0 = op(e4, e0)) | ~ sP0), inference(subsumption_resolution, [], [f799, f162])).
fof(f162, plain, ~ (e1 = e3), inference(cnf_transformation, [], [f1])).
fof(f799, plain, ((e1 = e3) | (e0 = op(e4, e0)) | ~ sP0), inference(forward_demodulation, [], [f798, f188])).
fof(f798, plain, ((e1 = op(e4, e1)) | (e0 = op(e4, e0)) | ~ sP0), inference(subsumption_resolution, [], [f797, f161])).
fof(f161, plain, ~ (e1 = e2), inference(cnf_transformation, [], [f1])).
fof(f797, plain, ((e1 = e2) | (e1 = op(e4, e1)) | (e0 = op(e4, e0)) | ~ sP0), inference(forward_demodulation, [], [f796, f189])).
fof(f796, plain, ((e2 = op(e4, e2)) | (e1 = op(e4, e1)) | (e0 = op(e4, e0)) | ~ sP0), inference(subsumption_resolution, [], [f795, f166])).
fof(f166, plain, ~ (e3 = e4), inference(cnf_transformation, [], [f1])).
fof(f795, plain, ((e3 = e4) | (e2 = op(e4, e2)) | (e1 = op(e4, e1)) | (e0 = op(e4, e0)) | ~ sP0), inference(forward_demodulation, [], [f794, f190])).
fof(f794, plain, ((e3 = op(e4, e3)) | (e2 = op(e4, e2)) | (e1 = op(e4, e1)) | (e0 = op(e4, e0)) | ~ sP0), inference(subsumption_resolution, [], [f793, f160])).
fof(f160, plain, ~ (e0 = e4), inference(cnf_transformation, [], [f1])).
fof(f793, plain, ((e0 = e4) | (e3 = op(e4, e3)) | (e2 = op(e4, e2)) | (e1 = op(e4, e1)) | (e0 = op(e4, e0)) | ~ sP0), inference(forward_demodulation, [], [f566, f191])).
fof(f566, plain, ((e4 = op(e4, e4)) | (e3 = op(e4, e3)) | (e2 = op(e4, e2)) | (e1 = op(e4, e1)) | (e0 = op(e4, e0)) | ~ sP0), inference(cnf_transformation, [], [f156])).
fof(f156, plain, ((((e4 = op(e4, e4)) | (e3 = op(e4, e3)) | (e2 = op(e4, e2)) | (e1 = op(e4, e1)) | (e0 = op(e4, e0))) & ((e4 = op(e3, e4)) | (e3 = op(e3, e3)) | (e2 = op(e3, e2)) | (e1 = op(e3, e1)) | (e0 = op(e3, e0))) & ((e4 = op(e2, e4)) | (e3 = op(e2, e3)) | (e2 = op(e2, e2)) | (e1 = op(e2, e1)) | (e0 = op(e2, e0))) & ((e4 = op(e1, e4)) | (e3 = op(e1, e3)) | (e2 = op(e1, e2)) | (e1 = op(e1, e1)) | (e0 = op(e1, e0))) & ((e4 = op(e0, e4)) | (e3 = op(e0, e3)) | (e2 = op(e0, e2)) | (e1 = op(e0, e1)) | (e0 = op(e0, e0)))) | ~ sP0), inference(nnf_transformation, [], [f6])).
fof(f6, plain, ((((e4 = op(e4, e4)) | (e3 = op(e4, e3)) | (e2 = op(e4, e2)) | (e1 = op(e4, e1)) | (e0 = op(e4, e0))) & ((e4 = op(e3, e4)) | (e3 = op(e3, e3)) | (e2 = op(e3, e2)) | (e1 = op(e3, e1)) | (e0 = op(e3, e0))) & ((e4 = op(e2, e4)) | (e3 = op(e2, e3)) | (e2 = op(e2, e2)) | (e1 = op(e2, e1)) | (e0 = op(e2, e0))) & ((e4 = op(e1, e4)) | (e3 = op(e1, e3)) | (e2 = op(e1, e2)) | (e1 = op(e1, e1)) | (e0 = op(e1, e0))) & ((e4 = op(e0, e4)) | (e3 = op(e0, e3)) | (e2 = op(e0, e2)) | (e1 = op(e0, e1)) | (e0 = op(e0, e0)))) | ~ sP0), inference(usedef, [], [e6])).
fof(e6, plain, (sP0 <=> (((e4 = op(e4, e4)) | (e3 = op(e4, e3)) | (e2 = op(e4, e2)) | (e1 = op(e4, e1)) | (e0 = op(e4, e0))) & ((e4 = op(e3, e4)) | (e3 = op(e3, e3)) | (e2 = op(e3, e2)) | (e1 = op(e3, e1)) | (e0 = op(e3, e0))) & ((e4 = op(e2, e4)) | (e3 = op(e2, e3)) | (e2 = op(e2, e2)) | (e1 = op(e2, e1)) | (e0 = op(e2, e0))) & ((e4 = op(e1, e4)) | (e3 = op(e1, e3)) | (e2 = op(e1, e2)) | (e1 = op(e1, e1)) | (e0 = op(e1, e0))) & ((e4 = op(e0, e4)) | (e3 = op(e0, e3)) | (e2 = op(e0, e2)) | (e1 = op(e0, e1)) | (e0 = op(e0, e0))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f1315, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP0), inference(subsumption_resolution, [], [f1314, f792])).
fof(f792, plain, ~ sP1, inference(subsumption_resolution, [], [f561, f167])).
fof(f561, plain, (~ (e4 = op(e0, e0)) | ~ sP1), inference(cnf_transformation, [], [f155])).
fof(f155, plain, ((~ (e4 = op(e0, e0)) & ~ (e3 = op(e0, e0)) & ~ (e2 = op(e0, e0)) & ~ (e1 = op(e0, e0)) & ~ (e0 = op(e0, e0))) | ~ sP1), inference(nnf_transformation, [], [f7])).
fof(f7, plain, ((~ (e4 = op(e0, e0)) & ~ (e3 = op(e0, e0)) & ~ (e2 = op(e0, e0)) & ~ (e1 = op(e0, e0)) & ~ (e0 = op(e0, e0))) | ~ sP1), inference(usedef, [], [e7])).
fof(e7, plain, (sP1 <=> (~ (e4 = op(e0, e0)) & ~ (e3 = op(e0, e0)) & ~ (e2 = op(e0, e0)) & ~ (e1 = op(e0, e0)) & ~ (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f1314, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP1 | sP0), inference(subsumption_resolution, [], [f1313, f791])).
fof(f791, plain, ~ sP2, inference(subsumption_resolution, [], [f554, f168])).
fof(f554, plain, (~ (e2 = op(e0, e1)) | ~ sP2), inference(cnf_transformation, [], [f154])).
fof(f154, plain, ((~ (e4 = op(e0, e1)) & ~ (e3 = op(e0, e1)) & ~ (e2 = op(e0, e1)) & ~ (e1 = op(e0, e1)) & ~ (e0 = op(e0, e1))) | ~ sP2), inference(nnf_transformation, [], [f8])).
fof(f8, plain, ((~ (e4 = op(e0, e1)) & ~ (e3 = op(e0, e1)) & ~ (e2 = op(e0, e1)) & ~ (e1 = op(e0, e1)) & ~ (e0 = op(e0, e1))) | ~ sP2), inference(usedef, [], [e8])).
fof(e8, plain, (sP2 <=> (~ (e4 = op(e0, e1)) & ~ (e3 = op(e0, e1)) & ~ (e2 = op(e0, e1)) & ~ (e1 = op(e0, e1)) & ~ (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f1313, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1312, f788])).
fof(f788, plain, ~ sP3, inference(subsumption_resolution, [], [f547, f169])).
fof(f547, plain, (~ (e0 = op(e0, e2)) | ~ sP3), inference(cnf_transformation, [], [f153])).
fof(f153, plain, ((~ (e4 = op(e0, e2)) & ~ (e3 = op(e0, e2)) & ~ (e2 = op(e0, e2)) & ~ (e1 = op(e0, e2)) & ~ (e0 = op(e0, e2))) | ~ sP3), inference(nnf_transformation, [], [f9])).
fof(f9, plain, ((~ (e4 = op(e0, e2)) & ~ (e3 = op(e0, e2)) & ~ (e2 = op(e0, e2)) & ~ (e1 = op(e0, e2)) & ~ (e0 = op(e0, e2))) | ~ sP3), inference(usedef, [], [e9])).
fof(e9, plain, (sP3 <=> (~ (e4 = op(e0, e2)) & ~ (e3 = op(e0, e2)) & ~ (e2 = op(e0, e2)) & ~ (e1 = op(e0, e2)) & ~ (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f1312, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1311, f783])).
fof(f783, plain, ~ sP4, inference(subsumption_resolution, [], [f543, f170])).
fof(f543, plain, (~ (e1 = op(e0, e3)) | ~ sP4), inference(cnf_transformation, [], [f152])).
fof(f152, plain, ((~ (e4 = op(e0, e3)) & ~ (e3 = op(e0, e3)) & ~ (e2 = op(e0, e3)) & ~ (e1 = op(e0, e3)) & ~ (e0 = op(e0, e3))) | ~ sP4), inference(nnf_transformation, [], [f10])).
fof(f10, plain, ((~ (e4 = op(e0, e3)) & ~ (e3 = op(e0, e3)) & ~ (e2 = op(e0, e3)) & ~ (e1 = op(e0, e3)) & ~ (e0 = op(e0, e3))) | ~ sP4), inference(usedef, [], [e10])).
fof(e10, plain, (sP4 <=> (~ (e4 = op(e0, e3)) & ~ (e3 = op(e0, e3)) & ~ (e2 = op(e0, e3)) & ~ (e1 = op(e0, e3)) & ~ (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f1311, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1310, f779])).
fof(f779, plain, ~ sP5, inference(subsumption_resolution, [], [f540, f171])).
fof(f540, plain, (~ (e3 = op(e0, e4)) | ~ sP5), inference(cnf_transformation, [], [f151])).
fof(f151, plain, ((~ (e4 = op(e0, e4)) & ~ (e3 = op(e0, e4)) & ~ (e2 = op(e0, e4)) & ~ (e1 = op(e0, e4)) & ~ (e0 = op(e0, e4))) | ~ sP5), inference(nnf_transformation, [], [f11])).
fof(f11, plain, ((~ (e4 = op(e0, e4)) & ~ (e3 = op(e0, e4)) & ~ (e2 = op(e0, e4)) & ~ (e1 = op(e0, e4)) & ~ (e0 = op(e0, e4))) | ~ sP5), inference(usedef, [], [e11])).
fof(e11, plain, (sP5 <=> (~ (e4 = op(e0, e4)) & ~ (e3 = op(e0, e4)) & ~ (e2 = op(e0, e4)) & ~ (e1 = op(e0, e4)) & ~ (e0 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f1310, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1309, f777])).
fof(f777, plain, ~ sP6, inference(subsumption_resolution, [], [f532, f172])).
fof(f532, plain, (~ (e0 = op(e1, e0)) | ~ sP6), inference(cnf_transformation, [], [f150])).
fof(f150, plain, ((~ (e4 = op(e1, e0)) & ~ (e3 = op(e1, e0)) & ~ (e2 = op(e1, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e1, e0))) | ~ sP6), inference(nnf_transformation, [], [f12])).
fof(f12, plain, ((~ (e4 = op(e1, e0)) & ~ (e3 = op(e1, e0)) & ~ (e2 = op(e1, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e1, e0))) | ~ sP6), inference(usedef, [], [e12])).
fof(e12, plain, (sP6 <=> (~ (e4 = op(e1, e0)) & ~ (e3 = op(e1, e0)) & ~ (e2 = op(e1, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f1309, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1308, f772])).
fof(f772, plain, ~ sP7, inference(subsumption_resolution, [], [f528, f173])).
fof(f528, plain, (~ (e1 = op(e1, e1)) | ~ sP7), inference(cnf_transformation, [], [f149])).
fof(f149, plain, ((~ (e4 = op(e1, e1)) & ~ (e3 = op(e1, e1)) & ~ (e2 = op(e1, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e1, e1))) | ~ sP7), inference(nnf_transformation, [], [f13])).
fof(f13, plain, ((~ (e4 = op(e1, e1)) & ~ (e3 = op(e1, e1)) & ~ (e2 = op(e1, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e1, e1))) | ~ sP7), inference(usedef, [], [e13])).
fof(e13, plain, (sP7 <=> (~ (e4 = op(e1, e1)) & ~ (e3 = op(e1, e1)) & ~ (e2 = op(e1, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f1308, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1307, f768])).
fof(f768, plain, ~ sP8, inference(subsumption_resolution, [], [f524, f174])).
fof(f524, plain, (~ (e2 = op(e1, e2)) | ~ sP8), inference(cnf_transformation, [], [f148])).
fof(f148, plain, ((~ (e4 = op(e1, e2)) & ~ (e3 = op(e1, e2)) & ~ (e2 = op(e1, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e1, e2))) | ~ sP8), inference(nnf_transformation, [], [f14])).
fof(f14, plain, ((~ (e4 = op(e1, e2)) & ~ (e3 = op(e1, e2)) & ~ (e2 = op(e1, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e1, e2))) | ~ sP8), inference(usedef, [], [e14])).
fof(e14, plain, (sP8 <=> (~ (e4 = op(e1, e2)) & ~ (e3 = op(e1, e2)) & ~ (e2 = op(e1, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f1307, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1306, f765])).
fof(f765, plain, ~ sP9, inference(subsumption_resolution, [], [f520, f175])).
fof(f520, plain, (~ (e3 = op(e1, e3)) | ~ sP9), inference(cnf_transformation, [], [f147])).
fof(f147, plain, ((~ (e4 = op(e1, e3)) & ~ (e3 = op(e1, e3)) & ~ (e2 = op(e1, e3)) & ~ (e1 = op(e1, e3)) & ~ (e0 = op(e1, e3))) | ~ sP9), inference(nnf_transformation, [], [f15])).
fof(f15, plain, ((~ (e4 = op(e1, e3)) & ~ (e3 = op(e1, e3)) & ~ (e2 = op(e1, e3)) & ~ (e1 = op(e1, e3)) & ~ (e0 = op(e1, e3))) | ~ sP9), inference(usedef, [], [e15])).
fof(e15, plain, (sP9 <=> (~ (e4 = op(e1, e3)) & ~ (e3 = op(e1, e3)) & ~ (e2 = op(e1, e3)) & ~ (e1 = op(e1, e3)) & ~ (e0 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f1306, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1305, f763])).
fof(f763, plain, ~ sP10, inference(subsumption_resolution, [], [f516, f176])).
fof(f516, plain, (~ (e4 = op(e1, e4)) | ~ sP10), inference(cnf_transformation, [], [f146])).
fof(f146, plain, ((~ (e4 = op(e1, e4)) & ~ (e3 = op(e1, e4)) & ~ (e2 = op(e1, e4)) & ~ (e1 = op(e1, e4)) & ~ (e0 = op(e1, e4))) | ~ sP10), inference(nnf_transformation, [], [f16])).
fof(f16, plain, ((~ (e4 = op(e1, e4)) & ~ (e3 = op(e1, e4)) & ~ (e2 = op(e1, e4)) & ~ (e1 = op(e1, e4)) & ~ (e0 = op(e1, e4))) | ~ sP10), inference(usedef, [], [e16])).
fof(e16, plain, (sP10 <=> (~ (e4 = op(e1, e4)) & ~ (e3 = op(e1, e4)) & ~ (e2 = op(e1, e4)) & ~ (e1 = op(e1, e4)) & ~ (e0 = op(e1, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f1305, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1304, f762])).
fof(f762, plain, ~ sP11, inference(subsumption_resolution, [], [f508, f177])).
fof(f508, plain, (~ (e1 = op(e2, e0)) | ~ sP11), inference(cnf_transformation, [], [f145])).
fof(f145, plain, ((~ (e4 = op(e2, e0)) & ~ (e3 = op(e2, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e2, e0)) & ~ (e0 = op(e2, e0))) | ~ sP11), inference(nnf_transformation, [], [f17])).
fof(f17, plain, ((~ (e4 = op(e2, e0)) & ~ (e3 = op(e2, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e2, e0)) & ~ (e0 = op(e2, e0))) | ~ sP11), inference(usedef, [], [e17])).
fof(e17, plain, (sP11 <=> (~ (e4 = op(e2, e0)) & ~ (e3 = op(e2, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e2, e0)) & ~ (e0 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f1304, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1303, f758])).
fof(f758, plain, ~ sP12, inference(subsumption_resolution, [], [f506, f178])).
fof(f506, plain, (~ (e4 = op(e2, e1)) | ~ sP12), inference(cnf_transformation, [], [f144])).
fof(f144, plain, ((~ (e4 = op(e2, e1)) & ~ (e3 = op(e2, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e2, e1)) & ~ (e0 = op(e2, e1))) | ~ sP12), inference(nnf_transformation, [], [f18])).
fof(f18, plain, ((~ (e4 = op(e2, e1)) & ~ (e3 = op(e2, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e2, e1)) & ~ (e0 = op(e2, e1))) | ~ sP12), inference(usedef, [], [e18])).
fof(e18, plain, (sP12 <=> (~ (e4 = op(e2, e1)) & ~ (e3 = op(e2, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e2, e1)) & ~ (e0 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f1303, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1302, f757])).
fof(f757, plain, ~ sP13, inference(subsumption_resolution, [], [f500, f179])).
fof(f500, plain, (~ (e3 = op(e2, e2)) | ~ sP13), inference(cnf_transformation, [], [f143])).
fof(f143, plain, ((~ (e4 = op(e2, e2)) & ~ (e3 = op(e2, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e2, e2)) & ~ (e0 = op(e2, e2))) | ~ sP13), inference(nnf_transformation, [], [f19])).
fof(f19, plain, ((~ (e4 = op(e2, e2)) & ~ (e3 = op(e2, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e2, e2)) & ~ (e0 = op(e2, e2))) | ~ sP13), inference(usedef, [], [e19])).
fof(e19, plain, (sP13 <=> (~ (e4 = op(e2, e2)) & ~ (e3 = op(e2, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e2, e2)) & ~ (e0 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f1302, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1301, f755])).
fof(f755, plain, ~ sP14, inference(subsumption_resolution, [], [f492, f180])).
fof(f492, plain, (~ (e0 = op(e2, e3)) | ~ sP14), inference(cnf_transformation, [], [f142])).
fof(f142, plain, ((~ (e4 = op(e2, e3)) & ~ (e3 = op(e2, e3)) & ~ (e2 = op(e2, e3)) & ~ (e1 = op(e2, e3)) & ~ (e0 = op(e2, e3))) | ~ sP14), inference(nnf_transformation, [], [f20])).
fof(f20, plain, ((~ (e4 = op(e2, e3)) & ~ (e3 = op(e2, e3)) & ~ (e2 = op(e2, e3)) & ~ (e1 = op(e2, e3)) & ~ (e0 = op(e2, e3))) | ~ sP14), inference(usedef, [], [e20])).
fof(e20, plain, (sP14 <=> (~ (e4 = op(e2, e3)) & ~ (e3 = op(e2, e3)) & ~ (e2 = op(e2, e3)) & ~ (e1 = op(e2, e3)) & ~ (e0 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f1301, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1300, f750])).
fof(f750, plain, ~ sP15, inference(subsumption_resolution, [], [f489, f181])).
fof(f489, plain, (~ (e2 = op(e2, e4)) | ~ sP15), inference(cnf_transformation, [], [f141])).
fof(f141, plain, ((~ (e4 = op(e2, e4)) & ~ (e3 = op(e2, e4)) & ~ (e2 = op(e2, e4)) & ~ (e1 = op(e2, e4)) & ~ (e0 = op(e2, e4))) | ~ sP15), inference(nnf_transformation, [], [f21])).
fof(f21, plain, ((~ (e4 = op(e2, e4)) & ~ (e3 = op(e2, e4)) & ~ (e2 = op(e2, e4)) & ~ (e1 = op(e2, e4)) & ~ (e0 = op(e2, e4))) | ~ sP15), inference(usedef, [], [e21])).
fof(e21, plain, (sP15 <=> (~ (e4 = op(e2, e4)) & ~ (e3 = op(e2, e4)) & ~ (e2 = op(e2, e4)) & ~ (e1 = op(e2, e4)) & ~ (e0 = op(e2, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f1300, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1299, f747])).
fof(f747, plain, ~ sP16, inference(subsumption_resolution, [], [f485, f182])).
fof(f485, plain, (~ (e3 = op(e3, e0)) | ~ sP16), inference(cnf_transformation, [], [f140])).
fof(f140, plain, ((~ (e4 = op(e3, e0)) & ~ (e3 = op(e3, e0)) & ~ (e2 = op(e3, e0)) & ~ (e1 = op(e3, e0)) & ~ (e0 = op(e3, e0))) | ~ sP16), inference(nnf_transformation, [], [f22])).
fof(f22, plain, ((~ (e4 = op(e3, e0)) & ~ (e3 = op(e3, e0)) & ~ (e2 = op(e3, e0)) & ~ (e1 = op(e3, e0)) & ~ (e0 = op(e3, e0))) | ~ sP16), inference(usedef, [], [e22])).
fof(e22, plain, (sP16 <=> (~ (e4 = op(e3, e0)) & ~ (e3 = op(e3, e0)) & ~ (e2 = op(e3, e0)) & ~ (e1 = op(e3, e0)) & ~ (e0 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f1299, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1298, f745])).
fof(f745, plain, ~ sP17, inference(subsumption_resolution, [], [f477, f183])).
fof(f477, plain, (~ (e0 = op(e3, e1)) | ~ sP17), inference(cnf_transformation, [], [f139])).
fof(f139, plain, ((~ (e4 = op(e3, e1)) & ~ (e3 = op(e3, e1)) & ~ (e2 = op(e3, e1)) & ~ (e1 = op(e3, e1)) & ~ (e0 = op(e3, e1))) | ~ sP17), inference(nnf_transformation, [], [f23])).
fof(f23, plain, ((~ (e4 = op(e3, e1)) & ~ (e3 = op(e3, e1)) & ~ (e2 = op(e3, e1)) & ~ (e1 = op(e3, e1)) & ~ (e0 = op(e3, e1))) | ~ sP17), inference(usedef, [], [e23])).
fof(e23, plain, (sP17 <=> (~ (e4 = op(e3, e1)) & ~ (e3 = op(e3, e1)) & ~ (e2 = op(e3, e1)) & ~ (e1 = op(e3, e1)) & ~ (e0 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f1298, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1297, f740])).
fof(f740, plain, ~ sP18, inference(subsumption_resolution, [], [f476, f184])).
fof(f476, plain, (~ (e4 = op(e3, e2)) | ~ sP18), inference(cnf_transformation, [], [f138])).
fof(f138, plain, ((~ (e4 = op(e3, e2)) & ~ (e3 = op(e3, e2)) & ~ (e2 = op(e3, e2)) & ~ (e1 = op(e3, e2)) & ~ (e0 = op(e3, e2))) | ~ sP18), inference(nnf_transformation, [], [f24])).
fof(f24, plain, ((~ (e4 = op(e3, e2)) & ~ (e3 = op(e3, e2)) & ~ (e2 = op(e3, e2)) & ~ (e1 = op(e3, e2)) & ~ (e0 = op(e3, e2))) | ~ sP18), inference(usedef, [], [e24])).
fof(e24, plain, (sP18 <=> (~ (e4 = op(e3, e2)) & ~ (e3 = op(e3, e2)) & ~ (e2 = op(e3, e2)) & ~ (e1 = op(e3, e2)) & ~ (e0 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f1297, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1296, f739])).
fof(f739, plain, ~ sP19, inference(subsumption_resolution, [], [f469, f185])).
fof(f469, plain, (~ (e2 = op(e3, e3)) | ~ sP19), inference(cnf_transformation, [], [f137])).
fof(f137, plain, ((~ (e4 = op(e3, e3)) & ~ (e3 = op(e3, e3)) & ~ (e2 = op(e3, e3)) & ~ (e1 = op(e3, e3)) & ~ (e0 = op(e3, e3))) | ~ sP19), inference(nnf_transformation, [], [f25])).
fof(f25, plain, ((~ (e4 = op(e3, e3)) & ~ (e3 = op(e3, e3)) & ~ (e2 = op(e3, e3)) & ~ (e1 = op(e3, e3)) & ~ (e0 = op(e3, e3))) | ~ sP19), inference(usedef, [], [e25])).
fof(e25, plain, (sP19 <=> (~ (e4 = op(e3, e3)) & ~ (e3 = op(e3, e3)) & ~ (e2 = op(e3, e3)) & ~ (e1 = op(e3, e3)) & ~ (e0 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP19])])).
fof(f1296, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1295, f736])).
fof(f736, plain, ~ sP20, inference(subsumption_resolution, [], [f463, f186])).
fof(f463, plain, (~ (e1 = op(e3, e4)) | ~ sP20), inference(cnf_transformation, [], [f136])).
fof(f136, plain, ((~ (e4 = op(e3, e4)) & ~ (e3 = op(e3, e4)) & ~ (e2 = op(e3, e4)) & ~ (e1 = op(e3, e4)) & ~ (e0 = op(e3, e4))) | ~ sP20), inference(nnf_transformation, [], [f26])).
fof(f26, plain, ((~ (e4 = op(e3, e4)) & ~ (e3 = op(e3, e4)) & ~ (e2 = op(e3, e4)) & ~ (e1 = op(e3, e4)) & ~ (e0 = op(e3, e4))) | ~ sP20), inference(usedef, [], [e26])).
fof(e26, plain, (sP20 <=> (~ (e4 = op(e3, e4)) & ~ (e3 = op(e3, e4)) & ~ (e2 = op(e3, e4)) & ~ (e1 = op(e3, e4)) & ~ (e0 = op(e3, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP20])])).
fof(f1295, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1294, f732])).
fof(f732, plain, ~ sP21, inference(subsumption_resolution, [], [f459, f187])).
fof(f459, plain, (~ (e2 = op(e4, e0)) | ~ sP21), inference(cnf_transformation, [], [f135])).
fof(f135, plain, ((~ (e4 = op(e4, e0)) & ~ (e3 = op(e4, e0)) & ~ (e2 = op(e4, e0)) & ~ (e1 = op(e4, e0)) & ~ (e0 = op(e4, e0))) | ~ sP21), inference(nnf_transformation, [], [f27])).
fof(f27, plain, ((~ (e4 = op(e4, e0)) & ~ (e3 = op(e4, e0)) & ~ (e2 = op(e4, e0)) & ~ (e1 = op(e4, e0)) & ~ (e0 = op(e4, e0))) | ~ sP21), inference(usedef, [], [e27])).
fof(e27, plain, (sP21 <=> (~ (e4 = op(e4, e0)) & ~ (e3 = op(e4, e0)) & ~ (e2 = op(e4, e0)) & ~ (e1 = op(e4, e0)) & ~ (e0 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP21])])).
fof(f1294, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1293, f729])).
fof(f729, plain, ~ sP22, inference(subsumption_resolution, [], [f455, f188])).
fof(f455, plain, (~ (e3 = op(e4, e1)) | ~ sP22), inference(cnf_transformation, [], [f134])).
fof(f134, plain, ((~ (e4 = op(e4, e1)) & ~ (e3 = op(e4, e1)) & ~ (e2 = op(e4, e1)) & ~ (e1 = op(e4, e1)) & ~ (e0 = op(e4, e1))) | ~ sP22), inference(nnf_transformation, [], [f28])).
fof(f28, plain, ((~ (e4 = op(e4, e1)) & ~ (e3 = op(e4, e1)) & ~ (e2 = op(e4, e1)) & ~ (e1 = op(e4, e1)) & ~ (e0 = op(e4, e1))) | ~ sP22), inference(usedef, [], [e28])).
fof(e28, plain, (sP22 <=> (~ (e4 = op(e4, e1)) & ~ (e3 = op(e4, e1)) & ~ (e2 = op(e4, e1)) & ~ (e1 = op(e4, e1)) & ~ (e0 = op(e4, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP22])])).
fof(f1293, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1292, f727])).
fof(f727, plain, ~ sP23, inference(subsumption_resolution, [], [f448, f189])).
fof(f448, plain, (~ (e1 = op(e4, e2)) | ~ sP23), inference(cnf_transformation, [], [f133])).
fof(f133, plain, ((~ (e4 = op(e4, e2)) & ~ (e3 = op(e4, e2)) & ~ (e2 = op(e4, e2)) & ~ (e1 = op(e4, e2)) & ~ (e0 = op(e4, e2))) | ~ sP23), inference(nnf_transformation, [], [f29])).
fof(f29, plain, ((~ (e4 = op(e4, e2)) & ~ (e3 = op(e4, e2)) & ~ (e2 = op(e4, e2)) & ~ (e1 = op(e4, e2)) & ~ (e0 = op(e4, e2))) | ~ sP23), inference(usedef, [], [e29])).
fof(e29, plain, (sP23 <=> (~ (e4 = op(e4, e2)) & ~ (e3 = op(e4, e2)) & ~ (e2 = op(e4, e2)) & ~ (e1 = op(e4, e2)) & ~ (e0 = op(e4, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP23])])).
fof(f1292, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1291, f723])).
fof(f723, plain, ~ sP24, inference(subsumption_resolution, [], [f446, f190])).
fof(f446, plain, (~ (e4 = op(e4, e3)) | ~ sP24), inference(cnf_transformation, [], [f132])).
fof(f132, plain, ((~ (e4 = op(e4, e3)) & ~ (e3 = op(e4, e3)) & ~ (e2 = op(e4, e3)) & ~ (e1 = op(e4, e3)) & ~ (e0 = op(e4, e3))) | ~ sP24), inference(nnf_transformation, [], [f30])).
fof(f30, plain, ((~ (e4 = op(e4, e3)) & ~ (e3 = op(e4, e3)) & ~ (e2 = op(e4, e3)) & ~ (e1 = op(e4, e3)) & ~ (e0 = op(e4, e3))) | ~ sP24), inference(usedef, [], [e30])).
fof(e30, plain, (sP24 <=> (~ (e4 = op(e4, e3)) & ~ (e3 = op(e4, e3)) & ~ (e2 = op(e4, e3)) & ~ (e1 = op(e4, e3)) & ~ (e0 = op(e4, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP24])])).
fof(f1291, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1290, f722])).
fof(f722, plain, ~ sP25, inference(subsumption_resolution, [], [f437, f191])).
fof(f437, plain, (~ (e0 = op(e4, e4)) | ~ sP25), inference(cnf_transformation, [], [f131])).
fof(f131, plain, ((~ (e4 = op(e4, e4)) & ~ (e3 = op(e4, e4)) & ~ (e2 = op(e4, e4)) & ~ (e1 = op(e4, e4)) & ~ (e0 = op(e4, e4))) | ~ sP25), inference(nnf_transformation, [], [f31])).
fof(f31, plain, ((~ (e4 = op(e4, e4)) & ~ (e3 = op(e4, e4)) & ~ (e2 = op(e4, e4)) & ~ (e1 = op(e4, e4)) & ~ (e0 = op(e4, e4))) | ~ sP25), inference(usedef, [], [e31])).
fof(e31, plain, (sP25 <=> (~ (e4 = op(e4, e4)) & ~ (e3 = op(e4, e4)) & ~ (e2 = op(e4, e4)) & ~ (e1 = op(e4, e4)) & ~ (e0 = op(e4, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP25])])).
fof(f1290, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1289, f717])).
fof(f717, plain, ~ sP26, inference(subsumption_resolution, [], [f434, f169])).
fof(f434, plain, (~ (e0 = op(e0, e2)) | ~ sP26), inference(cnf_transformation, [], [f130])).
fof(f130, plain, ((~ (e0 = op(e0, e4)) & ~ (e0 = op(e0, e3)) & ~ (e0 = op(e0, e2)) & ~ (e0 = op(e0, e1)) & ~ (e0 = op(e0, e0))) | ~ sP26), inference(nnf_transformation, [], [f32])).
fof(f32, plain, ((~ (e0 = op(e0, e4)) & ~ (e0 = op(e0, e3)) & ~ (e0 = op(e0, e2)) & ~ (e0 = op(e0, e1)) & ~ (e0 = op(e0, e0))) | ~ sP26), inference(usedef, [], [e32])).
fof(e32, plain, (sP26 <=> (~ (e0 = op(e0, e4)) & ~ (e0 = op(e0, e3)) & ~ (e0 = op(e0, e2)) & ~ (e0 = op(e0, e1)) & ~ (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP26])])).
fof(f1289, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1288, f714])).
fof(f714, plain, ~ sP27, inference(subsumption_resolution, [], [f428, f172])).
fof(f428, plain, (~ (e0 = op(e1, e0)) | ~ sP27), inference(cnf_transformation, [], [f129])).
fof(f129, plain, ((~ (e0 = op(e4, e0)) & ~ (e0 = op(e3, e0)) & ~ (e0 = op(e2, e0)) & ~ (e0 = op(e1, e0)) & ~ (e0 = op(e0, e0))) | ~ sP27), inference(nnf_transformation, [], [f33])).
fof(f33, plain, ((~ (e0 = op(e4, e0)) & ~ (e0 = op(e3, e0)) & ~ (e0 = op(e2, e0)) & ~ (e0 = op(e1, e0)) & ~ (e0 = op(e0, e0))) | ~ sP27), inference(usedef, [], [e33])).
fof(e33, plain, (sP27 <=> (~ (e0 = op(e4, e0)) & ~ (e0 = op(e3, e0)) & ~ (e0 = op(e2, e0)) & ~ (e0 = op(e1, e0)) & ~ (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP27])])).
fof(f1288, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1287, f710])).
fof(f710, plain, ~ sP28, inference(subsumption_resolution, [], [f425, f170])).
fof(f425, plain, (~ (e1 = op(e0, e3)) | ~ sP28), inference(cnf_transformation, [], [f128])).
fof(f128, plain, ((~ (e1 = op(e0, e4)) & ~ (e1 = op(e0, e3)) & ~ (e1 = op(e0, e2)) & ~ (e1 = op(e0, e1)) & ~ (e1 = op(e0, e0))) | ~ sP28), inference(nnf_transformation, [], [f34])).
fof(f34, plain, ((~ (e1 = op(e0, e4)) & ~ (e1 = op(e0, e3)) & ~ (e1 = op(e0, e2)) & ~ (e1 = op(e0, e1)) & ~ (e1 = op(e0, e0))) | ~ sP28), inference(usedef, [], [e34])).
fof(e34, plain, (sP28 <=> (~ (e1 = op(e0, e4)) & ~ (e1 = op(e0, e3)) & ~ (e1 = op(e0, e2)) & ~ (e1 = op(e0, e1)) & ~ (e1 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP28])])).
fof(f1287, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1286, f708])).
fof(f708, plain, ~ sP29, inference(subsumption_resolution, [], [f419, f177])).
fof(f419, plain, (~ (e1 = op(e2, e0)) | ~ sP29), inference(cnf_transformation, [], [f127])).
fof(f127, plain, ((~ (e1 = op(e4, e0)) & ~ (e1 = op(e3, e0)) & ~ (e1 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e1 = op(e0, e0))) | ~ sP29), inference(nnf_transformation, [], [f35])).
fof(f35, plain, ((~ (e1 = op(e4, e0)) & ~ (e1 = op(e3, e0)) & ~ (e1 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e1 = op(e0, e0))) | ~ sP29), inference(usedef, [], [e35])).
fof(e35, plain, (sP29 <=> (~ (e1 = op(e4, e0)) & ~ (e1 = op(e3, e0)) & ~ (e1 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e1 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP29])])).
fof(f1286, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1285, f705])).
fof(f705, plain, ~ sP30, inference(subsumption_resolution, [], [f413, f168])).
fof(f413, plain, (~ (e2 = op(e0, e1)) | ~ sP30), inference(cnf_transformation, [], [f126])).
fof(f126, plain, ((~ (e2 = op(e0, e4)) & ~ (e2 = op(e0, e3)) & ~ (e2 = op(e0, e2)) & ~ (e2 = op(e0, e1)) & ~ (e2 = op(e0, e0))) | ~ sP30), inference(nnf_transformation, [], [f36])).
fof(f36, plain, ((~ (e2 = op(e0, e4)) & ~ (e2 = op(e0, e3)) & ~ (e2 = op(e0, e2)) & ~ (e2 = op(e0, e1)) & ~ (e2 = op(e0, e0))) | ~ sP30), inference(usedef, [], [e36])).
fof(e36, plain, (sP30 <=> (~ (e2 = op(e0, e4)) & ~ (e2 = op(e0, e3)) & ~ (e2 = op(e0, e2)) & ~ (e2 = op(e0, e1)) & ~ (e2 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP30])])).
fof(f1285, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1284, f701])).
fof(f701, plain, ~ sP31, inference(subsumption_resolution, [], [f411, f187])).
fof(f411, plain, (~ (e2 = op(e4, e0)) | ~ sP31), inference(cnf_transformation, [], [f125])).
fof(f125, plain, ((~ (e2 = op(e4, e0)) & ~ (e2 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e2 = op(e1, e0)) & ~ (e2 = op(e0, e0))) | ~ sP31), inference(nnf_transformation, [], [f37])).
fof(f37, plain, ((~ (e2 = op(e4, e0)) & ~ (e2 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e2 = op(e1, e0)) & ~ (e2 = op(e0, e0))) | ~ sP31), inference(usedef, [], [e37])).
fof(e37, plain, (sP31 <=> (~ (e2 = op(e4, e0)) & ~ (e2 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e2 = op(e1, e0)) & ~ (e2 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP31])])).
fof(f1284, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1283, f700])).
fof(f700, plain, ~ sP32, inference(subsumption_resolution, [], [f406, f171])).
fof(f406, plain, (~ (e3 = op(e0, e4)) | ~ sP32), inference(cnf_transformation, [], [f124])).
fof(f124, plain, ((~ (e3 = op(e0, e4)) & ~ (e3 = op(e0, e3)) & ~ (e3 = op(e0, e2)) & ~ (e3 = op(e0, e1)) & ~ (e3 = op(e0, e0))) | ~ sP32), inference(nnf_transformation, [], [f38])).
fof(f38, plain, ((~ (e3 = op(e0, e4)) & ~ (e3 = op(e0, e3)) & ~ (e3 = op(e0, e2)) & ~ (e3 = op(e0, e1)) & ~ (e3 = op(e0, e0))) | ~ sP32), inference(usedef, [], [e38])).
fof(e38, plain, (sP32 <=> (~ (e3 = op(e0, e4)) & ~ (e3 = op(e0, e3)) & ~ (e3 = op(e0, e2)) & ~ (e3 = op(e0, e1)) & ~ (e3 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP32])])).
fof(f1283, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1282, f699])).
fof(f699, plain, ~ sP33, inference(subsumption_resolution, [], [f400, f182])).
fof(f400, plain, (~ (e3 = op(e3, e0)) | ~ sP33), inference(cnf_transformation, [], [f123])).
fof(f123, plain, ((~ (e3 = op(e4, e0)) & ~ (e3 = op(e3, e0)) & ~ (e3 = op(e2, e0)) & ~ (e3 = op(e1, e0)) & ~ (e3 = op(e0, e0))) | ~ sP33), inference(nnf_transformation, [], [f39])).
fof(f39, plain, ((~ (e3 = op(e4, e0)) & ~ (e3 = op(e3, e0)) & ~ (e3 = op(e2, e0)) & ~ (e3 = op(e1, e0)) & ~ (e3 = op(e0, e0))) | ~ sP33), inference(usedef, [], [e39])).
fof(e39, plain, (sP33 <=> (~ (e3 = op(e4, e0)) & ~ (e3 = op(e3, e0)) & ~ (e3 = op(e2, e0)) & ~ (e3 = op(e1, e0)) & ~ (e3 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP33])])).
fof(f1282, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1281, f697])).
fof(f697, plain, ~ sP34, inference(subsumption_resolution, [], [f392, f167])).
fof(f392, plain, (~ (e4 = op(e0, e0)) | ~ sP34), inference(cnf_transformation, [], [f122])).
fof(f122, plain, ((~ (e4 = op(e0, e4)) & ~ (e4 = op(e0, e3)) & ~ (e4 = op(e0, e2)) & ~ (e4 = op(e0, e1)) & ~ (e4 = op(e0, e0))) | ~ sP34), inference(nnf_transformation, [], [f40])).
fof(f40, plain, ((~ (e4 = op(e0, e4)) & ~ (e4 = op(e0, e3)) & ~ (e4 = op(e0, e2)) & ~ (e4 = op(e0, e1)) & ~ (e4 = op(e0, e0))) | ~ sP34), inference(usedef, [], [e40])).
fof(e40, plain, (sP34 <=> (~ (e4 = op(e0, e4)) & ~ (e4 = op(e0, e3)) & ~ (e4 = op(e0, e2)) & ~ (e4 = op(e0, e1)) & ~ (e4 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP34])])).
fof(f1281, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1280, f692])).
fof(f692, plain, ~ sP35, inference(subsumption_resolution, [], [f387, f167])).
fof(f387, plain, (~ (e4 = op(e0, e0)) | ~ sP35), inference(cnf_transformation, [], [f121])).
fof(f121, plain, ((~ (e4 = op(e4, e0)) & ~ (e4 = op(e3, e0)) & ~ (e4 = op(e2, e0)) & ~ (e4 = op(e1, e0)) & ~ (e4 = op(e0, e0))) | ~ sP35), inference(nnf_transformation, [], [f41])).
fof(f41, plain, ((~ (e4 = op(e4, e0)) & ~ (e4 = op(e3, e0)) & ~ (e4 = op(e2, e0)) & ~ (e4 = op(e1, e0)) & ~ (e4 = op(e0, e0))) | ~ sP35), inference(usedef, [], [e41])).
fof(e41, plain, (sP35 <=> (~ (e4 = op(e4, e0)) & ~ (e4 = op(e3, e0)) & ~ (e4 = op(e2, e0)) & ~ (e4 = op(e1, e0)) & ~ (e4 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP35])])).
fof(f1280, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1279, f687])).
fof(f687, plain, ~ sP36, inference(subsumption_resolution, [], [f382, f172])).
fof(f382, plain, (~ (e0 = op(e1, e0)) | ~ sP36), inference(cnf_transformation, [], [f120])).
fof(f120, plain, ((~ (e0 = op(e1, e4)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e1, e0))) | ~ sP36), inference(nnf_transformation, [], [f42])).
fof(f42, plain, ((~ (e0 = op(e1, e4)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e1, e0))) | ~ sP36), inference(usedef, [], [e42])).
fof(e42, plain, (sP36 <=> (~ (e0 = op(e1, e4)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP36])])).
fof(f1279, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1278, f682])).
fof(f682, plain, ~ sP37, inference(subsumption_resolution, [], [f380, f183])).
fof(f380, plain, (~ (e0 = op(e3, e1)) | ~ sP37), inference(cnf_transformation, [], [f119])).
fof(f119, plain, ((~ (e0 = op(e4, e1)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e0, e1))) | ~ sP37), inference(nnf_transformation, [], [f43])).
fof(f43, plain, ((~ (e0 = op(e4, e1)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e0, e1))) | ~ sP37), inference(usedef, [], [e43])).
fof(e43, plain, (sP37 <=> (~ (e0 = op(e4, e1)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP37])])).
fof(f1278, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1277, f680])).
fof(f680, plain, ~ sP38, inference(subsumption_resolution, [], [f373, f173])).
fof(f373, plain, (~ (e1 = op(e1, e1)) | ~ sP38), inference(cnf_transformation, [], [f118])).
fof(f118, plain, ((~ (e1 = op(e1, e4)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e1, e0))) | ~ sP38), inference(nnf_transformation, [], [f44])).
fof(f44, plain, ((~ (e1 = op(e1, e4)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e1, e0))) | ~ sP38), inference(usedef, [], [e44])).
fof(e44, plain, (sP38 <=> (~ (e1 = op(e1, e4)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP38])])).
fof(f1277, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1276, f676])).
fof(f676, plain, ~ sP39, inference(subsumption_resolution, [], [f368, f173])).
fof(f368, plain, (~ (e1 = op(e1, e1)) | ~ sP39), inference(cnf_transformation, [], [f117])).
fof(f117, plain, ((~ (e1 = op(e4, e1)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e0, e1))) | ~ sP39), inference(nnf_transformation, [], [f45])).
fof(f45, plain, ((~ (e1 = op(e4, e1)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e0, e1))) | ~ sP39), inference(usedef, [], [e45])).
fof(e45, plain, (sP39 <=> (~ (e1 = op(e4, e1)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP39])])).
fof(f1276, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1275, f672])).
fof(f672, plain, ~ sP40, inference(subsumption_resolution, [], [f364, f174])).
fof(f364, plain, (~ (e2 = op(e1, e2)) | ~ sP40), inference(cnf_transformation, [], [f116])).
fof(f116, plain, ((~ (e2 = op(e1, e4)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e1, e0))) | ~ sP40), inference(nnf_transformation, [], [f46])).
fof(f46, plain, ((~ (e2 = op(e1, e4)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e1, e0))) | ~ sP40), inference(usedef, [], [e46])).
fof(e46, plain, (sP40 <=> (~ (e2 = op(e1, e4)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP40])])).
fof(f1275, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1274, f669])).
fof(f669, plain, ~ sP41, inference(subsumption_resolution, [], [f357, f168])).
fof(f357, plain, (~ (e2 = op(e0, e1)) | ~ sP41), inference(cnf_transformation, [], [f115])).
fof(f115, plain, ((~ (e2 = op(e4, e1)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e0, e1))) | ~ sP41), inference(nnf_transformation, [], [f47])).
fof(f47, plain, ((~ (e2 = op(e4, e1)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e0, e1))) | ~ sP41), inference(usedef, [], [e47])).
fof(e47, plain, (sP41 <=> (~ (e2 = op(e4, e1)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP41])])).
fof(f1274, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1273, f664])).
fof(f664, plain, ~ sP42, inference(subsumption_resolution, [], [f355, f175])).
fof(f355, plain, (~ (e3 = op(e1, e3)) | ~ sP42), inference(cnf_transformation, [], [f114])).
fof(f114, plain, ((~ (e3 = op(e1, e4)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e1, e0))) | ~ sP42), inference(nnf_transformation, [], [f48])).
fof(f48, plain, ((~ (e3 = op(e1, e4)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e1, e0))) | ~ sP42), inference(usedef, [], [e48])).
fof(e48, plain, (sP42 <=> (~ (e3 = op(e1, e4)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP42])])).
fof(f1273, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1272, f662])).
fof(f662, plain, ~ sP43, inference(subsumption_resolution, [], [f351, f188])).
fof(f351, plain, (~ (e3 = op(e4, e1)) | ~ sP43), inference(cnf_transformation, [], [f113])).
fof(f113, plain, ((~ (e3 = op(e4, e1)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e0, e1))) | ~ sP43), inference(nnf_transformation, [], [f49])).
fof(f49, plain, ((~ (e3 = op(e4, e1)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e0, e1))) | ~ sP43), inference(usedef, [], [e49])).
fof(e49, plain, (sP43 <=> (~ (e3 = op(e4, e1)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP43])])).
fof(f1272, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1271, f661])).
fof(f661, plain, ~ sP44, inference(subsumption_resolution, [], [f346, f176])).
fof(f346, plain, (~ (e4 = op(e1, e4)) | ~ sP44), inference(cnf_transformation, [], [f112])).
fof(f112, plain, ((~ (e4 = op(e1, e4)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e1, e0))) | ~ sP44), inference(nnf_transformation, [], [f50])).
fof(f50, plain, ((~ (e4 = op(e1, e4)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e1, e0))) | ~ sP44), inference(usedef, [], [e50])).
fof(e50, plain, (sP44 <=> (~ (e4 = op(e1, e4)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP44])])).
fof(f1271, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1270, f660])).
fof(f660, plain, ~ sP45, inference(subsumption_resolution, [], [f339, f178])).
fof(f339, plain, (~ (e4 = op(e2, e1)) | ~ sP45), inference(cnf_transformation, [], [f111])).
fof(f111, plain, ((~ (e4 = op(e4, e1)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e0, e1))) | ~ sP45), inference(nnf_transformation, [], [f51])).
fof(f51, plain, ((~ (e4 = op(e4, e1)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e0, e1))) | ~ sP45), inference(usedef, [], [e51])).
fof(e51, plain, (sP45 <=> (~ (e4 = op(e4, e1)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP45])])).
fof(f1270, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1269, f657])).
fof(f657, plain, ~ sP46, inference(subsumption_resolution, [], [f335, f180])).
fof(f335, plain, (~ (e0 = op(e2, e3)) | ~ sP46), inference(cnf_transformation, [], [f110])).
fof(f110, plain, ((~ (e0 = op(e2, e4)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e2, e0))) | ~ sP46), inference(nnf_transformation, [], [f52])).
fof(f52, plain, ((~ (e0 = op(e2, e4)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e2, e0))) | ~ sP46), inference(usedef, [], [e52])).
fof(e52, plain, (sP46 <=> (~ (e0 = op(e2, e4)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP46])])).
fof(f1269, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1268, f655])).
fof(f655, plain, ~ sP47, inference(subsumption_resolution, [], [f327, f169])).
fof(f327, plain, (~ (e0 = op(e0, e2)) | ~ sP47), inference(cnf_transformation, [], [f109])).
fof(f109, plain, ((~ (e0 = op(e4, e2)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e0, e2))) | ~ sP47), inference(nnf_transformation, [], [f53])).
fof(f53, plain, ((~ (e0 = op(e4, e2)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e0, e2))) | ~ sP47), inference(usedef, [], [e53])).
fof(e53, plain, (sP47 <=> (~ (e0 = op(e4, e2)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP47])])).
fof(f1268, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1267, f650])).
fof(f650, plain, ~ sP48, inference(subsumption_resolution, [], [f322, f177])).
fof(f322, plain, (~ (e1 = op(e2, e0)) | ~ sP48), inference(cnf_transformation, [], [f108])).
fof(f108, plain, ((~ (e1 = op(e2, e4)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e2, e0))) | ~ sP48), inference(nnf_transformation, [], [f54])).
fof(f54, plain, ((~ (e1 = op(e2, e4)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e2, e0))) | ~ sP48), inference(usedef, [], [e54])).
fof(e54, plain, (sP48 <=> (~ (e1 = op(e2, e4)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP48])])).
fof(f1267, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1266, f645])).
fof(f645, plain, ~ sP49, inference(subsumption_resolution, [], [f321, f189])).
fof(f321, plain, (~ (e1 = op(e4, e2)) | ~ sP49), inference(cnf_transformation, [], [f107])).
fof(f107, plain, ((~ (e1 = op(e4, e2)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e0, e2))) | ~ sP49), inference(nnf_transformation, [], [f55])).
fof(f55, plain, ((~ (e1 = op(e4, e2)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e0, e2))) | ~ sP49), inference(usedef, [], [e55])).
fof(e55, plain, (sP49 <=> (~ (e1 = op(e4, e2)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP49])])).
fof(f1266, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1265, f644])).
fof(f644, plain, ~ sP50, inference(subsumption_resolution, [], [f316, f181])).
fof(f316, plain, (~ (e2 = op(e2, e4)) | ~ sP50), inference(cnf_transformation, [], [f106])).
fof(f106, plain, ((~ (e2 = op(e2, e4)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e2, e0))) | ~ sP50), inference(nnf_transformation, [], [f56])).
fof(f56, plain, ((~ (e2 = op(e2, e4)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e2, e0))) | ~ sP50), inference(usedef, [], [e56])).
fof(e56, plain, (sP50 <=> (~ (e2 = op(e2, e4)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP50])])).
fof(f1265, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1264, f643])).
fof(f643, plain, ~ sP51, inference(subsumption_resolution, [], [f308, f174])).
fof(f308, plain, (~ (e2 = op(e1, e2)) | ~ sP51), inference(cnf_transformation, [], [f105])).
fof(f105, plain, ((~ (e2 = op(e4, e2)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e0, e2))) | ~ sP51), inference(nnf_transformation, [], [f57])).
fof(f57, plain, ((~ (e2 = op(e4, e2)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e0, e2))) | ~ sP51), inference(usedef, [], [e57])).
fof(e57, plain, (sP51 <=> (~ (e2 = op(e4, e2)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP51])])).
fof(f1264, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1263, f639])).
fof(f639, plain, ~ sP52, inference(subsumption_resolution, [], [f304, f179])).
fof(f304, plain, (~ (e3 = op(e2, e2)) | ~ sP52), inference(cnf_transformation, [], [f104])).
fof(f104, plain, ((~ (e3 = op(e2, e4)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e2, e0))) | ~ sP52), inference(nnf_transformation, [], [f58])).
fof(f58, plain, ((~ (e3 = op(e2, e4)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e2, e0))) | ~ sP52), inference(usedef, [], [e58])).
fof(e58, plain, (sP52 <=> (~ (e3 = op(e2, e4)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP52])])).
fof(f1263, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1262, f636])).
fof(f636, plain, ~ sP53, inference(subsumption_resolution, [], [f299, f179])).
fof(f299, plain, (~ (e3 = op(e2, e2)) | ~ sP53), inference(cnf_transformation, [], [f103])).
fof(f103, plain, ((~ (e3 = op(e4, e2)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e0, e2))) | ~ sP53), inference(nnf_transformation, [], [f59])).
fof(f59, plain, ((~ (e3 = op(e4, e2)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e0, e2))) | ~ sP53), inference(usedef, [], [e59])).
fof(e59, plain, (sP53 <=> (~ (e3 = op(e4, e2)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP53])])).
fof(f1262, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1261, f633])).
fof(f633, plain, ~ sP54, inference(subsumption_resolution, [], [f293, f178])).
fof(f293, plain, (~ (e4 = op(e2, e1)) | ~ sP54), inference(cnf_transformation, [], [f102])).
fof(f102, plain, ((~ (e4 = op(e2, e4)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e2, e0))) | ~ sP54), inference(nnf_transformation, [], [f60])).
fof(f60, plain, ((~ (e4 = op(e2, e4)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e2, e0))) | ~ sP54), inference(usedef, [], [e60])).
fof(e60, plain, (sP54 <=> (~ (e4 = op(e2, e4)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP54])])).
fof(f1261, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1260, f629])).
fof(f629, plain, ~ sP55, inference(subsumption_resolution, [], [f290, f184])).
fof(f290, plain, (~ (e4 = op(e3, e2)) | ~ sP55), inference(cnf_transformation, [], [f101])).
fof(f101, plain, ((~ (e4 = op(e4, e2)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e0, e2))) | ~ sP55), inference(nnf_transformation, [], [f61])).
fof(f61, plain, ((~ (e4 = op(e4, e2)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e0, e2))) | ~ sP55), inference(usedef, [], [e61])).
fof(e61, plain, (sP55 <=> (~ (e4 = op(e4, e2)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP55])])).
fof(f1260, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1259, f627])).
fof(f627, plain, ~ sP56, inference(subsumption_resolution, [], [f283, f183])).
fof(f283, plain, (~ (e0 = op(e3, e1)) | ~ sP56), inference(cnf_transformation, [], [f100])).
fof(f100, plain, ((~ (e0 = op(e3, e4)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e3, e0))) | ~ sP56), inference(nnf_transformation, [], [f62])).
fof(f62, plain, ((~ (e0 = op(e3, e4)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e3, e0))) | ~ sP56), inference(usedef, [], [e62])).
fof(e62, plain, (sP56 <=> (~ (e0 = op(e3, e4)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP56])])).
fof(f1259, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1258, f623])).
fof(f623, plain, ~ sP57, inference(subsumption_resolution, [], [f279, f180])).
fof(f279, plain, (~ (e0 = op(e2, e3)) | ~ sP57), inference(cnf_transformation, [], [f99])).
fof(f99, plain, ((~ (e0 = op(e4, e3)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e0, e3))) | ~ sP57), inference(nnf_transformation, [], [f63])).
fof(f63, plain, ((~ (e0 = op(e4, e3)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e0, e3))) | ~ sP57), inference(usedef, [], [e63])).
fof(e63, plain, (sP57 <=> (~ (e0 = op(e4, e3)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP57])])).
fof(f1258, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1257, f620])).
fof(f620, plain, ~ sP58, inference(subsumption_resolution, [], [f276, f186])).
fof(f276, plain, (~ (e1 = op(e3, e4)) | ~ sP58), inference(cnf_transformation, [], [f98])).
fof(f98, plain, ((~ (e1 = op(e3, e4)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e3, e0))) | ~ sP58), inference(nnf_transformation, [], [f64])).
fof(f64, plain, ((~ (e1 = op(e3, e4)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e3, e0))) | ~ sP58), inference(usedef, [], [e64])).
fof(e64, plain, (sP58 <=> (~ (e1 = op(e3, e4)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP58])])).
fof(f1257, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1256, f619])).
fof(f619, plain, ~ sP59, inference(subsumption_resolution, [], [f267, f170])).
fof(f267, plain, (~ (e1 = op(e0, e3)) | ~ sP59), inference(cnf_transformation, [], [f97])).
fof(f97, plain, ((~ (e1 = op(e4, e3)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e0, e3))) | ~ sP59), inference(nnf_transformation, [], [f65])).
fof(f65, plain, ((~ (e1 = op(e4, e3)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e0, e3))) | ~ sP59), inference(usedef, [], [e65])).
fof(e65, plain, (sP59 <=> (~ (e1 = op(e4, e3)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP59])])).
fof(f1256, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1255, f614])).
fof(f614, plain, ~ sP60, inference(subsumption_resolution, [], [f265, f185])).
fof(f265, plain, (~ (e2 = op(e3, e3)) | ~ sP60), inference(cnf_transformation, [], [f96])).
fof(f96, plain, ((~ (e2 = op(e3, e4)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e3, e0))) | ~ sP60), inference(nnf_transformation, [], [f66])).
fof(f66, plain, ((~ (e2 = op(e3, e4)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e3, e0))) | ~ sP60), inference(usedef, [], [e66])).
fof(e66, plain, (sP60 <=> (~ (e2 = op(e3, e4)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP60])])).
fof(f1255, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1254, f612])).
fof(f612, plain, ~ sP61, inference(subsumption_resolution, [], [f260, f185])).
fof(f260, plain, (~ (e2 = op(e3, e3)) | ~ sP61), inference(cnf_transformation, [], [f95])).
fof(f95, plain, ((~ (e2 = op(e4, e3)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e0, e3))) | ~ sP61), inference(nnf_transformation, [], [f67])).
fof(f67, plain, ((~ (e2 = op(e4, e3)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e0, e3))) | ~ sP61), inference(usedef, [], [e67])).
fof(e67, plain, (sP61 <=> (~ (e2 = op(e4, e3)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP61])])).
fof(f1254, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1253, f610])).
fof(f610, plain, ~ sP62, inference(subsumption_resolution, [], [f252, f182])).
fof(f252, plain, (~ (e3 = op(e3, e0)) | ~ sP62), inference(cnf_transformation, [], [f94])).
fof(f94, plain, ((~ (e3 = op(e3, e4)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e3, e0))) | ~ sP62), inference(nnf_transformation, [], [f68])).
fof(f68, plain, ((~ (e3 = op(e3, e4)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e3, e0))) | ~ sP62), inference(usedef, [], [e68])).
fof(e68, plain, (sP62 <=> (~ (e3 = op(e3, e4)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP62])])).
fof(f1253, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1252, f605])).
fof(f605, plain, ~ sP63, inference(subsumption_resolution, [], [f248, f175])).
fof(f248, plain, (~ (e3 = op(e1, e3)) | ~ sP63), inference(cnf_transformation, [], [f93])).
fof(f93, plain, ((~ (e3 = op(e4, e3)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e0, e3))) | ~ sP63), inference(nnf_transformation, [], [f69])).
fof(f69, plain, ((~ (e3 = op(e4, e3)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e0, e3))) | ~ sP63), inference(usedef, [], [e69])).
fof(e69, plain, (sP63 <=> (~ (e3 = op(e4, e3)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP63])])).
fof(f1252, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1251, f601])).
fof(f601, plain, ~ sP64, inference(subsumption_resolution, [], [f244, f184])).
fof(f244, plain, (~ (e4 = op(e3, e2)) | ~ sP64), inference(cnf_transformation, [], [f92])).
fof(f92, plain, ((~ (e4 = op(e3, e4)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e3, e0))) | ~ sP64), inference(nnf_transformation, [], [f70])).
fof(f70, plain, ((~ (e4 = op(e3, e4)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e3, e0))) | ~ sP64), inference(usedef, [], [e70])).
fof(e70, plain, (sP64 <=> (~ (e4 = op(e3, e4)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP64])])).
fof(f1251, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1250, f598])).
fof(f598, plain, ~ sP65, inference(subsumption_resolution, [], [f241, f190])).
fof(f241, plain, (~ (e4 = op(e4, e3)) | ~ sP65), inference(cnf_transformation, [], [f91])).
fof(f91, plain, ((~ (e4 = op(e4, e3)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e0, e3))) | ~ sP65), inference(nnf_transformation, [], [f71])).
fof(f71, plain, ((~ (e4 = op(e4, e3)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e0, e3))) | ~ sP65), inference(usedef, [], [e71])).
fof(e71, plain, (sP65 <=> (~ (e4 = op(e4, e3)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP65])])).
fof(f1250, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1249, f597])).
fof(f597, plain, ~ sP66, inference(subsumption_resolution, [], [f236, f191])).
fof(f236, plain, (~ (e0 = op(e4, e4)) | ~ sP66), inference(cnf_transformation, [], [f90])).
fof(f90, plain, ((~ (e0 = op(e4, e4)) & ~ (e0 = op(e4, e3)) & ~ (e0 = op(e4, e2)) & ~ (e0 = op(e4, e1)) & ~ (e0 = op(e4, e0))) | ~ sP66), inference(nnf_transformation, [], [f72])).
fof(f72, plain, ((~ (e0 = op(e4, e4)) & ~ (e0 = op(e4, e3)) & ~ (e0 = op(e4, e2)) & ~ (e0 = op(e4, e1)) & ~ (e0 = op(e4, e0))) | ~ sP66), inference(usedef, [], [e72])).
fof(e72, plain, (sP66 <=> (~ (e0 = op(e4, e4)) & ~ (e0 = op(e4, e3)) & ~ (e0 = op(e4, e2)) & ~ (e0 = op(e4, e1)) & ~ (e0 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP66])])).
fof(f1249, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1248, f596])).
fof(f596, plain, ~ sP67, inference(subsumption_resolution, [], [f231, f191])).
fof(f231, plain, (~ (e0 = op(e4, e4)) | ~ sP67), inference(cnf_transformation, [], [f89])).
fof(f89, plain, ((~ (e0 = op(e4, e4)) & ~ (e0 = op(e3, e4)) & ~ (e0 = op(e2, e4)) & ~ (e0 = op(e1, e4)) & ~ (e0 = op(e0, e4))) | ~ sP67), inference(nnf_transformation, [], [f73])).
fof(f73, plain, ((~ (e0 = op(e4, e4)) & ~ (e0 = op(e3, e4)) & ~ (e0 = op(e2, e4)) & ~ (e0 = op(e1, e4)) & ~ (e0 = op(e0, e4))) | ~ sP67), inference(usedef, [], [e73])).
fof(e73, plain, (sP67 <=> (~ (e0 = op(e4, e4)) & ~ (e0 = op(e3, e4)) & ~ (e0 = op(e2, e4)) & ~ (e0 = op(e1, e4)) & ~ (e0 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP67])])).
fof(f1248, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1247, f595])).
fof(f595, plain, ~ sP68, inference(subsumption_resolution, [], [f224, f189])).
fof(f224, plain, (~ (e1 = op(e4, e2)) | ~ sP68), inference(cnf_transformation, [], [f88])).
fof(f88, plain, ((~ (e1 = op(e4, e4)) & ~ (e1 = op(e4, e3)) & ~ (e1 = op(e4, e2)) & ~ (e1 = op(e4, e1)) & ~ (e1 = op(e4, e0))) | ~ sP68), inference(nnf_transformation, [], [f74])).
fof(f74, plain, ((~ (e1 = op(e4, e4)) & ~ (e1 = op(e4, e3)) & ~ (e1 = op(e4, e2)) & ~ (e1 = op(e4, e1)) & ~ (e1 = op(e4, e0))) | ~ sP68), inference(usedef, [], [e74])).
fof(e74, plain, (sP68 <=> (~ (e1 = op(e4, e4)) & ~ (e1 = op(e4, e3)) & ~ (e1 = op(e4, e2)) & ~ (e1 = op(e4, e1)) & ~ (e1 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP68])])).
fof(f1247, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1246, f592])).
fof(f592, plain, ~ sP69, inference(subsumption_resolution, [], [f220, f186])).
fof(f220, plain, (~ (e1 = op(e3, e4)) | ~ sP69), inference(cnf_transformation, [], [f87])).
fof(f87, plain, ((~ (e1 = op(e4, e4)) & ~ (e1 = op(e3, e4)) & ~ (e1 = op(e2, e4)) & ~ (e1 = op(e1, e4)) & ~ (e1 = op(e0, e4))) | ~ sP69), inference(nnf_transformation, [], [f75])).
fof(f75, plain, ((~ (e1 = op(e4, e4)) & ~ (e1 = op(e3, e4)) & ~ (e1 = op(e2, e4)) & ~ (e1 = op(e1, e4)) & ~ (e1 = op(e0, e4))) | ~ sP69), inference(usedef, [], [e75])).
fof(e75, plain, (sP69 <=> (~ (e1 = op(e4, e4)) & ~ (e1 = op(e3, e4)) & ~ (e1 = op(e2, e4)) & ~ (e1 = op(e1, e4)) & ~ (e1 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP69])])).
fof(f1246, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1245, f590])).
fof(f590, plain, ~ sP70, inference(subsumption_resolution, [], [f212, f187])).
fof(f212, plain, (~ (e2 = op(e4, e0)) | ~ sP70), inference(cnf_transformation, [], [f86])).
fof(f86, plain, ((~ (e2 = op(e4, e4)) & ~ (e2 = op(e4, e3)) & ~ (e2 = op(e4, e2)) & ~ (e2 = op(e4, e1)) & ~ (e2 = op(e4, e0))) | ~ sP70), inference(nnf_transformation, [], [f76])).
fof(f76, plain, ((~ (e2 = op(e4, e4)) & ~ (e2 = op(e4, e3)) & ~ (e2 = op(e4, e2)) & ~ (e2 = op(e4, e1)) & ~ (e2 = op(e4, e0))) | ~ sP70), inference(usedef, [], [e76])).
fof(e76, plain, (sP70 <=> (~ (e2 = op(e4, e4)) & ~ (e2 = op(e4, e3)) & ~ (e2 = op(e4, e2)) & ~ (e2 = op(e4, e1)) & ~ (e2 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP70])])).
fof(f1245, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1244, f585])).
fof(f585, plain, ~ sP71, inference(subsumption_resolution, [], [f209, f181])).
fof(f209, plain, (~ (e2 = op(e2, e4)) | ~ sP71), inference(cnf_transformation, [], [f85])).
fof(f85, plain, ((~ (e2 = op(e4, e4)) & ~ (e2 = op(e3, e4)) & ~ (e2 = op(e2, e4)) & ~ (e2 = op(e1, e4)) & ~ (e2 = op(e0, e4))) | ~ sP71), inference(nnf_transformation, [], [f77])).
fof(f77, plain, ((~ (e2 = op(e4, e4)) & ~ (e2 = op(e3, e4)) & ~ (e2 = op(e2, e4)) & ~ (e2 = op(e1, e4)) & ~ (e2 = op(e0, e4))) | ~ sP71), inference(usedef, [], [e77])).
fof(e77, plain, (sP71 <=> (~ (e2 = op(e4, e4)) & ~ (e2 = op(e3, e4)) & ~ (e2 = op(e2, e4)) & ~ (e2 = op(e1, e4)) & ~ (e2 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP71])])).
fof(f1244, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1243, f582])).
fof(f582, plain, ~ sP72, inference(subsumption_resolution, [], [f203, f188])).
fof(f203, plain, (~ (e3 = op(e4, e1)) | ~ sP72), inference(cnf_transformation, [], [f84])).
fof(f84, plain, ((~ (e3 = op(e4, e4)) & ~ (e3 = op(e4, e3)) & ~ (e3 = op(e4, e2)) & ~ (e3 = op(e4, e1)) & ~ (e3 = op(e4, e0))) | ~ sP72), inference(nnf_transformation, [], [f78])).
fof(f78, plain, ((~ (e3 = op(e4, e4)) & ~ (e3 = op(e4, e3)) & ~ (e3 = op(e4, e2)) & ~ (e3 = op(e4, e1)) & ~ (e3 = op(e4, e0))) | ~ sP72), inference(usedef, [], [e78])).
fof(e78, plain, (sP72 <=> (~ (e3 = op(e4, e4)) & ~ (e3 = op(e4, e3)) & ~ (e3 = op(e4, e2)) & ~ (e3 = op(e4, e1)) & ~ (e3 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP72])])).
fof(f1243, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1242, f578])).
fof(f578, plain, ~ sP73, inference(subsumption_resolution, [], [f197, f171])).
fof(f197, plain, (~ (e3 = op(e0, e4)) | ~ sP73), inference(cnf_transformation, [], [f83])).
fof(f83, plain, ((~ (e3 = op(e4, e4)) & ~ (e3 = op(e3, e4)) & ~ (e3 = op(e2, e4)) & ~ (e3 = op(e1, e4)) & ~ (e3 = op(e0, e4))) | ~ sP73), inference(nnf_transformation, [], [f79])).
fof(f79, plain, ((~ (e3 = op(e4, e4)) & ~ (e3 = op(e3, e4)) & ~ (e3 = op(e2, e4)) & ~ (e3 = op(e1, e4)) & ~ (e3 = op(e0, e4))) | ~ sP73), inference(usedef, [], [e79])).
fof(e79, plain, (sP73 <=> (~ (e3 = op(e4, e4)) & ~ (e3 = op(e3, e4)) & ~ (e3 = op(e2, e4)) & ~ (e3 = op(e1, e4)) & ~ (e3 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP73])])).
fof(f1242, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1241, f573])).
fof(f573, plain, ~ sP74, inference(subsumption_resolution, [], [f195, f190])).
fof(f195, plain, (~ (e4 = op(e4, e3)) | ~ sP74), inference(cnf_transformation, [], [f82])).
fof(f82, plain, ((~ (e4 = op(e4, e4)) & ~ (e4 = op(e4, e3)) & ~ (e4 = op(e4, e2)) & ~ (e4 = op(e4, e1)) & ~ (e4 = op(e4, e0))) | ~ sP74), inference(nnf_transformation, [], [f80])).
fof(f80, plain, ((~ (e4 = op(e4, e4)) & ~ (e4 = op(e4, e3)) & ~ (e4 = op(e4, e2)) & ~ (e4 = op(e4, e1)) & ~ (e4 = op(e4, e0))) | ~ sP74), inference(usedef, [], [e80])).
fof(e80, plain, (sP74 <=> (~ (e4 = op(e4, e4)) & ~ (e4 = op(e4, e3)) & ~ (e4 = op(e4, e2)) & ~ (e4 = op(e4, e1)) & ~ (e4 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP74])])).
fof(f1241, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f568, f176])).
fof(f568, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | ~ (e4 = op(e1, e4)) | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f81])).
fof(f81, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | (~ (e4 = op(e4, e4)) & ~ (e4 = op(e3, e4)) & ~ (e4 = op(e2, e4)) & ~ (e4 = op(e1, e4)) & ~ (e4 = op(e0, e4))) | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(definition_folding, [], [f5, e80, e79, e78, e77, e76, e75, e74, e73, e72, e71, e70, e69, e68, e67, e66, e65, e64, e63, e62, e61, e60, e59, e58, e57, e56, e55, e54, e53, e52, e51, e50, e49, e48, e47, e46, e45, e44, e43, e42, e41, e40, e39, e38, e37, e36, e35, e34, e33, e32, e31, e30, e29, e28, e27, e26, e25, e24, e23, e22, e21, e20, e19, e18, e17, e16, e15, e14, e13, e12, e11, e10, e9, e8, e7, e6])).
fof(f5, plain, (~ (e4 = op(op(e4, e4), op(e4, e4))) | ~ (e4 = op(op(e3, e4), op(e4, e3))) | ~ (e4 = op(op(e2, e4), op(e4, e2))) | ~ (e4 = op(op(e1, e4), op(e4, e1))) | ~ (e4 = op(op(e0, e4), op(e4, e0))) | ~ (e3 = op(op(e4, e3), op(e3, e4))) | ~ (e3 = op(op(e3, e3), op(e3, e3))) | ~ (e3 = op(op(e2, e3), op(e3, e2))) | ~ (e3 = op(op(e1, e3), op(e3, e1))) | ~ (e3 = op(op(e0, e3), op(e3, e0))) | ~ (e2 = op(op(e4, e2), op(e2, e4))) | ~ (e2 = op(op(e3, e2), op(e2, e3))) | ~ (e2 = op(op(e2, e2), op(e2, e2))) | ~ (e2 = op(op(e1, e2), op(e2, e1))) | ~ (e2 = op(op(e0, e2), op(e2, e0))) | ~ (e1 = op(op(e4, e1), op(e1, e4))) | ~ (e1 = op(op(e3, e1), op(e1, e3))) | ~ (e1 = op(op(e2, e1), op(e1, e2))) | ~ (e1 = op(op(e1, e1), op(e1, e1))) | ~ (e1 = op(op(e0, e1), op(e1, e0))) | ~ (e0 = op(op(e4, e0), op(e0, e4))) | ~ (e0 = op(op(e3, e0), op(e0, e3))) | ~ (e0 = op(op(e2, e0), op(e0, e2))) | ~ (e0 = op(op(e1, e0), op(e0, e1))) | ~ (e0 = op(op(e0, e0), op(e0, e0))) | (~ (e4 = op(e4, e4)) & ~ (e4 = op(e3, e4)) & ~ (e4 = op(e2, e4)) & ~ (e4 = op(e1, e4)) & ~ (e4 = op(e0, e4))) | (~ (e4 = op(e4, e4)) & ~ (e4 = op(e4, e3)) & ~ (e4 = op(e4, e2)) & ~ (e4 = op(e4, e1)) & ~ (e4 = op(e4, e0))) | (~ (e3 = op(e4, e4)) & ~ (e3 = op(e3, e4)) & ~ (e3 = op(e2, e4)) & ~ (e3 = op(e1, e4)) & ~ (e3 = op(e0, e4))) | (~ (e3 = op(e4, e4)) & ~ (e3 = op(e4, e3)) & ~ (e3 = op(e4, e2)) & ~ (e3 = op(e4, e1)) & ~ (e3 = op(e4, e0))) | (~ (e2 = op(e4, e4)) & ~ (e2 = op(e3, e4)) & ~ (e2 = op(e2, e4)) & ~ (e2 = op(e1, e4)) & ~ (e2 = op(e0, e4))) | (~ (e2 = op(e4, e4)) & ~ (e2 = op(e4, e3)) & ~ (e2 = op(e4, e2)) & ~ (e2 = op(e4, e1)) & ~ (e2 = op(e4, e0))) | (~ (e1 = op(e4, e4)) & ~ (e1 = op(e3, e4)) & ~ (e1 = op(e2, e4)) & ~ (e1 = op(e1, e4)) & ~ (e1 = op(e0, e4))) | (~ (e1 = op(e4, e4)) & ~ (e1 = op(e4, e3)) & ~ (e1 = op(e4, e2)) & ~ (e1 = op(e4, e1)) & ~ (e1 = op(e4, e0))) | (~ (e0 = op(e4, e4)) & ~ (e0 = op(e3, e4)) & ~ (e0 = op(e2, e4)) & ~ (e0 = op(e1, e4)) & ~ (e0 = op(e0, e4))) | (~ (e0 = op(e4, e4)) & ~ (e0 = op(e4, e3)) & ~ (e0 = op(e4, e2)) & ~ (e0 = op(e4, e1)) & ~ (e0 = op(e4, e0))) | (~ (e4 = op(e4, e3)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e0, e3))) | (~ (e4 = op(e3, e4)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e3, e0))) | (~ (e3 = op(e4, e3)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e0, e3))) | (~ (e3 = op(e3, e4)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e3, e0))) | (~ (e2 = op(e4, e3)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e0, e3))) | (~ (e2 = op(e3, e4)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e3, e0))) | (~ (e1 = op(e4, e3)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e0, e3))) | (~ (e1 = op(e3, e4)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e3, e0))) | (~ (e0 = op(e4, e3)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e0, e3))) | (~ (e0 = op(e3, e4)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e3, e0))) | (~ (e4 = op(e4, e2)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e0, e2))) | (~ (e4 = op(e2, e4)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e2, e0))) | (~ (e3 = op(e4, e2)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e0, e2))) | (~ (e3 = op(e2, e4)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e2, e0))) | (~ (e2 = op(e4, e2)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e0, e2))) | (~ (e2 = op(e2, e4)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e2, e0))) | (~ (e1 = op(e4, e2)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e0, e2))) | (~ (e1 = op(e2, e4)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e2, e0))) | (~ (e0 = op(e4, e2)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e0, e2))) | (~ (e0 = op(e2, e4)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e2, e0))) | (~ (e4 = op(e4, e1)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e0, e1))) | (~ (e4 = op(e1, e4)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e1, e0))) | (~ (e3 = op(e4, e1)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e0, e1))) | (~ (e3 = op(e1, e4)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e1, e0))) | (~ (e2 = op(e4, e1)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e0, e1))) | (~ (e2 = op(e1, e4)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e1, e0))) | (~ (e1 = op(e4, e1)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e0, e1))) | (~ (e1 = op(e1, e4)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e1, e0))) | (~ (e0 = op(e4, e1)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e0, e1))) | (~ (e0 = op(e1, e4)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e1, e0))) | (~ (e4 = op(e4, e0)) & ~ (e4 = op(e3, e0)) & ~ (e4 = op(e2, e0)) & ~ (e4 = op(e1, e0)) & ~ (e4 = op(e0, e0))) | (~ (e4 = op(e0, e4)) & ~ (e4 = op(e0, e3)) & ~ (e4 = op(e0, e2)) & ~ (e4 = op(e0, e1)) & ~ (e4 = op(e0, e0))) | (~ (e3 = op(e4, e0)) & ~ (e3 = op(e3, e0)) & ~ (e3 = op(e2, e0)) & ~ (e3 = op(e1, e0)) & ~ (e3 = op(e0, e0))) | (~ (e3 = op(e0, e4)) & ~ (e3 = op(e0, e3)) & ~ (e3 = op(e0, e2)) & ~ (e3 = op(e0, e1)) & ~ (e3 = op(e0, e0))) | (~ (e2 = op(e4, e0)) & ~ (e2 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e2 = op(e1, e0)) & ~ (e2 = op(e0, e0))) | (~ (e2 = op(e0, e4)) & ~ (e2 = op(e0, e3)) & ~ (e2 = op(e0, e2)) & ~ (e2 = op(e0, e1)) & ~ (e2 = op(e0, e0))) | (~ (e1 = op(e4, e0)) & ~ (e1 = op(e3, e0)) & ~ (e1 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e1 = op(e0, e0))) | (~ (e1 = op(e0, e4)) & ~ (e1 = op(e0, e3)) & ~ (e1 = op(e0, e2)) & ~ (e1 = op(e0, e1)) & ~ (e1 = op(e0, e0))) | (~ (e0 = op(e4, e0)) & ~ (e0 = op(e3, e0)) & ~ (e0 = op(e2, e0)) & ~ (e0 = op(e1, e0)) & ~ (e0 = op(e0, e0))) | (~ (e0 = op(e0, e4)) & ~ (e0 = op(e0, e3)) & ~ (e0 = op(e0, e2)) & ~ (e0 = op(e0, e1)) & ~ (e0 = op(e0, e0))) | (~ (e4 = op(e4, e4)) & ~ (e3 = op(e4, e4)) & ~ (e2 = op(e4, e4)) & ~ (e1 = op(e4, e4)) & ~ (e0 = op(e4, e4))) | (~ (e4 = op(e4, e3)) & ~ (e3 = op(e4, e3)) & ~ (e2 = op(e4, e3)) & ~ (e1 = op(e4, e3)) & ~ (e0 = op(e4, e3))) | (~ (e4 = op(e4, e2)) & ~ (e3 = op(e4, e2)) & ~ (e2 = op(e4, e2)) & ~ (e1 = op(e4, e2)) & ~ (e0 = op(e4, e2))) | (~ (e4 = op(e4, e1)) & ~ (e3 = op(e4, e1)) & ~ (e2 = op(e4, e1)) & ~ (e1 = op(e4, e1)) & ~ (e0 = op(e4, e1))) | (~ (e4 = op(e4, e0)) & ~ (e3 = op(e4, e0)) & ~ (e2 = op(e4, e0)) & ~ (e1 = op(e4, e0)) & ~ (e0 = op(e4, e0))) | (~ (e4 = op(e3, e4)) & ~ (e3 = op(e3, e4)) & ~ (e2 = op(e3, e4)) & ~ (e1 = op(e3, e4)) & ~ (e0 = op(e3, e4))) | (~ (e4 = op(e3, e3)) & ~ (e3 = op(e3, e3)) & ~ (e2 = op(e3, e3)) & ~ (e1 = op(e3, e3)) & ~ (e0 = op(e3, e3))) | (~ (e4 = op(e3, e2)) & ~ (e3 = op(e3, e2)) & ~ (e2 = op(e3, e2)) & ~ (e1 = op(e3, e2)) & ~ (e0 = op(e3, e2))) | (~ (e4 = op(e3, e1)) & ~ (e3 = op(e3, e1)) & ~ (e2 = op(e3, e1)) & ~ (e1 = op(e3, e1)) & ~ (e0 = op(e3, e1))) | (~ (e4 = op(e3, e0)) & ~ (e3 = op(e3, e0)) & ~ (e2 = op(e3, e0)) & ~ (e1 = op(e3, e0)) & ~ (e0 = op(e3, e0))) | (~ (e4 = op(e2, e4)) & ~ (e3 = op(e2, e4)) & ~ (e2 = op(e2, e4)) & ~ (e1 = op(e2, e4)) & ~ (e0 = op(e2, e4))) | (~ (e4 = op(e2, e3)) & ~ (e3 = op(e2, e3)) & ~ (e2 = op(e2, e3)) & ~ (e1 = op(e2, e3)) & ~ (e0 = op(e2, e3))) | (~ (e4 = op(e2, e2)) & ~ (e3 = op(e2, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e2, e2)) & ~ (e0 = op(e2, e2))) | (~ (e4 = op(e2, e1)) & ~ (e3 = op(e2, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e2, e1)) & ~ (e0 = op(e2, e1))) | (~ (e4 = op(e2, e0)) & ~ (e3 = op(e2, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e2, e0)) & ~ (e0 = op(e2, e0))) | (~ (e4 = op(e1, e4)) & ~ (e3 = op(e1, e4)) & ~ (e2 = op(e1, e4)) & ~ (e1 = op(e1, e4)) & ~ (e0 = op(e1, e4))) | (~ (e4 = op(e1, e3)) & ~ (e3 = op(e1, e3)) & ~ (e2 = op(e1, e3)) & ~ (e1 = op(e1, e3)) & ~ (e0 = op(e1, e3))) | (~ (e4 = op(e1, e2)) & ~ (e3 = op(e1, e2)) & ~ (e2 = op(e1, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e1, e2))) | (~ (e4 = op(e1, e1)) & ~ (e3 = op(e1, e1)) & ~ (e2 = op(e1, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e1, e1))) | (~ (e4 = op(e1, e0)) & ~ (e3 = op(e1, e0)) & ~ (e2 = op(e1, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e1, e0))) | (~ (e4 = op(e0, e4)) & ~ (e3 = op(e0, e4)) & ~ (e2 = op(e0, e4)) & ~ (e1 = op(e0, e4)) & ~ (e0 = op(e0, e4))) | (~ (e4 = op(e0, e3)) & ~ (e3 = op(e0, e3)) & ~ (e2 = op(e0, e3)) & ~ (e1 = op(e0, e3)) & ~ (e0 = op(e0, e3))) | (~ (e4 = op(e0, e2)) & ~ (e3 = op(e0, e2)) & ~ (e2 = op(e0, e2)) & ~ (e1 = op(e0, e2)) & ~ (e0 = op(e0, e2))) | (~ (e4 = op(e0, e1)) & ~ (e3 = op(e0, e1)) & ~ (e2 = op(e0, e1)) & ~ (e1 = op(e0, e1)) & ~ (e0 = op(e0, e1))) | (~ (e4 = op(e0, e0)) & ~ (e3 = op(e0, e0)) & ~ (e2 = op(e0, e0)) & ~ (e1 = op(e0, e0)) & ~ (e0 = op(e0, e0))) | (((e4 = op(e4, e4)) | (e3 = op(e4, e3)) | (e2 = op(e4, e2)) | (e1 = op(e4, e1)) | (e0 = op(e4, e0))) & ((e4 = op(e3, e4)) | (e3 = op(e3, e3)) | (e2 = op(e3, e2)) | (e1 = op(e3, e1)) | (e0 = op(e3, e0))) & ((e4 = op(e2, e4)) | (e3 = op(e2, e3)) | (e2 = op(e2, e2)) | (e1 = op(e2, e1)) | (e0 = op(e2, e0))) & ((e4 = op(e1, e4)) | (e3 = op(e1, e3)) | (e2 = op(e1, e2)) | (e1 = op(e1, e1)) | (e0 = op(e1, e0))) & ((e4 = op(e0, e4)) | (e3 = op(e0, e3)) | (e2 = op(e0, e2)) | (e1 = op(e0, e1)) | (e0 = op(e0, e0))))), inference(ennf_transformation, [], [f4])).
fof(f4, plain, ~ ((e4 = op(op(e4, e4), op(e4, e4))) & (e4 = op(op(e3, e4), op(e4, e3))) & (e4 = op(op(e2, e4), op(e4, e2))) & (e4 = op(op(e1, e4), op(e4, e1))) & (e4 = op(op(e0, e4), op(e4, e0))) & (e3 = op(op(e4, e3), op(e3, e4))) & (e3 = op(op(e3, e3), op(e3, e3))) & (e3 = op(op(e2, e3), op(e3, e2))) & (e3 = op(op(e1, e3), op(e3, e1))) & (e3 = op(op(e0, e3), op(e3, e0))) & (e2 = op(op(e4, e2), op(e2, e4))) & (e2 = op(op(e3, e2), op(e2, e3))) & (e2 = op(op(e2, e2), op(e2, e2))) & (e2 = op(op(e1, e2), op(e2, e1))) & (e2 = op(op(e0, e2), op(e2, e0))) & (e1 = op(op(e4, e1), op(e1, e4))) & (e1 = op(op(e3, e1), op(e1, e3))) & (e1 = op(op(e2, e1), op(e1, e2))) & (e1 = op(op(e1, e1), op(e1, e1))) & (e1 = op(op(e0, e1), op(e1, e0))) & (e0 = op(op(e4, e0), op(e0, e4))) & (e0 = op(op(e3, e0), op(e0, e3))) & (e0 = op(op(e2, e0), op(e0, e2))) & (e0 = op(op(e1, e0), op(e0, e1))) & (e0 = op(op(e0, e0), op(e0, e0))) & ((e4 = op(e4, e4)) | (e4 = op(e3, e4)) | (e4 = op(e2, e4)) | (e4 = op(e1, e4)) | (e4 = op(e0, e4))) & ((e4 = op(e4, e4)) | (e4 = op(e4, e3)) | (e4 = op(e4, e2)) | (e4 = op(e4, e1)) | (e4 = op(e4, e0))) & ((e3 = op(e4, e4)) | (e3 = op(e3, e4)) | (e3 = op(e2, e4)) | (e3 = op(e1, e4)) | (e3 = op(e0, e4))) & ((e3 = op(e4, e4)) | (e3 = op(e4, e3)) | (e3 = op(e4, e2)) | (e3 = op(e4, e1)) | (e3 = op(e4, e0))) & ((e2 = op(e4, e4)) | (e2 = op(e3, e4)) | (e2 = op(e2, e4)) | (e2 = op(e1, e4)) | (e2 = op(e0, e4))) & ((e2 = op(e4, e4)) | (e2 = op(e4, e3)) | (e2 = op(e4, e2)) | (e2 = op(e4, e1)) | (e2 = op(e4, e0))) & ((e1 = op(e4, e4)) | (e1 = op(e3, e4)) | (e1 = op(e2, e4)) | (e1 = op(e1, e4)) | (e1 = op(e0, e4))) & ((e1 = op(e4, e4)) | (e1 = op(e4, e3)) | (e1 = op(e4, e2)) | (e1 = op(e4, e1)) | (e1 = op(e4, e0))) & ((e0 = op(e4, e4)) | (e0 = op(e3, e4)) | (e0 = op(e2, e4)) | (e0 = op(e1, e4)) | (e0 = op(e0, e4))) & ((e0 = op(e4, e4)) | (e0 = op(e4, e3)) | (e0 = op(e4, e2)) | (e0 = op(e4, e1)) | (e0 = op(e4, e0))) & ((e4 = op(e4, e3)) | (e4 = op(e3, e3)) | (e4 = op(e2, e3)) | (e4 = op(e1, e3)) | (e4 = op(e0, e3))) & ((e4 = op(e3, e4)) | (e4 = op(e3, e3)) | (e4 = op(e3, e2)) | (e4 = op(e3, e1)) | (e4 = op(e3, e0))) & ((e3 = op(e4, e3)) | (e3 = op(e3, e3)) | (e3 = op(e2, e3)) | (e3 = op(e1, e3)) | (e3 = op(e0, e3))) & ((e3 = op(e3, e4)) | (e3 = op(e3, e3)) | (e3 = op(e3, e2)) | (e3 = op(e3, e1)) | (e3 = op(e3, e0))) & ((e2 = op(e4, e3)) | (e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(e0, e3))) & ((e2 = op(e3, e4)) | (e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, e0))) & ((e1 = op(e4, e3)) | (e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(e0, e3))) & ((e1 = op(e3, e4)) | (e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, e0))) & ((e0 = op(e4, e3)) | (e0 = op(e3, e3)) | (e0 = op(e2, e3)) | (e0 = op(e1, e3)) | (e0 = op(e0, e3))) & ((e0 = op(e3, e4)) | (e0 = op(e3, e3)) | (e0 = op(e3, e2)) | (e0 = op(e3, e1)) | (e0 = op(e3, e0))) & ((e4 = op(e4, e2)) | (e4 = op(e3, e2)) | (e4 = op(e2, e2)) | (e4 = op(e1, e2)) | (e4 = op(e0, e2))) & ((e4 = op(e2, e4)) | (e4 = op(e2, e3)) | (e4 = op(e2, e2)) | (e4 = op(e2, e1)) | (e4 = op(e2, e0))) & ((e3 = op(e4, e2)) | (e3 = op(e3, e2)) | (e3 = op(e2, e2)) | (e3 = op(e1, e2)) | (e3 = op(e0, e2))) & ((e3 = op(e2, e4)) | (e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, e0))) & ((e2 = op(e4, e2)) | (e2 = op(e3, e2)) | (e2 = op(e2, e2)) | (e2 = op(e1, e2)) | (e2 = op(e0, e2))) & ((e2 = op(e2, e4)) | (e2 = op(e2, e3)) | (e2 = op(e2, e2)) | (e2 = op(e2, e1)) | (e2 = op(e2, e0))) & ((e1 = op(e4, e2)) | (e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(e0, e2))) & ((e1 = op(e2, e4)) | (e1 = op(e2, e3)) | (e1 = op(e2, e2)) | (e1 = op(e2, e1)) | (e1 = op(e2, e0))) & ((e0 = op(e4, e2)) | (e0 = op(e3, e2)) | (e0 = op(e2, e2)) | (e0 = op(e1, e2)) | (e0 = op(e0, e2))) & ((e0 = op(e2, e4)) | (e0 = op(e2, e3)) | (e0 = op(e2, e2)) | (e0 = op(e2, e1)) | (e0 = op(e2, e0))) & ((e4 = op(e4, e1)) | (e4 = op(e3, e1)) | (e4 = op(e2, e1)) | (e4 = op(e1, e1)) | (e4 = op(e0, e1))) & ((e4 = op(e1, e4)) | (e4 = op(e1, e3)) | (e4 = op(e1, e2)) | (e4 = op(e1, e1)) | (e4 = op(e1, e0))) & ((e3 = op(e4, e1)) | (e3 = op(e3, e1)) | (e3 = op(e2, e1)) | (e3 = op(e1, e1)) | (e3 = op(e0, e1))) & ((e3 = op(e1, e4)) | (e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, e0))) & ((e2 = op(e4, e1)) | (e2 = op(e3, e1)) | (e2 = op(e2, e1)) | (e2 = op(e1, e1)) | (e2 = op(e0, e1))) & ((e2 = op(e1, e4)) | (e2 = op(e1, e3)) | (e2 = op(e1, e2)) | (e2 = op(e1, e1)) | (e2 = op(e1, e0))) & ((e1 = op(e4, e1)) | (e1 = op(e3, e1)) | (e1 = op(e2, e1)) | (e1 = op(e1, e1)) | (e1 = op(e0, e1))) & ((e1 = op(e1, e4)) | (e1 = op(e1, e3)) | (e1 = op(e1, e2)) | (e1 = op(e1, e1)) | (e1 = op(e1, e0))) & ((e0 = op(e4, e1)) | (e0 = op(e3, e1)) | (e0 = op(e2, e1)) | (e0 = op(e1, e1)) | (e0 = op(e0, e1))) & ((e0 = op(e1, e4)) | (e0 = op(e1, e3)) | (e0 = op(e1, e2)) | (e0 = op(e1, e1)) | (e0 = op(e1, e0))) & ((e4 = op(e4, e0)) | (e4 = op(e3, e0)) | (e4 = op(e2, e0)) | (e4 = op(e1, e0)) | (e4 = op(e0, e0))) & ((e4 = op(e0, e4)) | (e4 = op(e0, e3)) | (e4 = op(e0, e2)) | (e4 = op(e0, e1)) | (e4 = op(e0, e0))) & ((e3 = op(e4, e0)) | (e3 = op(e3, e0)) | (e3 = op(e2, e0)) | (e3 = op(e1, e0)) | (e3 = op(e0, e0))) & ((e3 = op(e0, e4)) | (e3 = op(e0, e3)) | (e3 = op(e0, e2)) | (e3 = op(e0, e1)) | (e3 = op(e0, e0))) & ((e2 = op(e4, e0)) | (e2 = op(e3, e0)) | (e2 = op(e2, e0)) | (e2 = op(e1, e0)) | (e2 = op(e0, e0))) & ((e2 = op(e0, e4)) | (e2 = op(e0, e3)) | (e2 = op(e0, e2)) | (e2 = op(e0, e1)) | (e2 = op(e0, e0))) & ((e1 = op(e4, e0)) | (e1 = op(e3, e0)) | (e1 = op(e2, e0)) | (e1 = op(e1, e0)) | (e1 = op(e0, e0))) & ((e1 = op(e0, e4)) | (e1 = op(e0, e3)) | (e1 = op(e0, e2)) | (e1 = op(e0, e1)) | (e1 = op(e0, e0))) & ((e0 = op(e4, e0)) | (e0 = op(e3, e0)) | (e0 = op(e2, e0)) | (e0 = op(e1, e0)) | (e0 = op(e0, e0))) & ((e0 = op(e0, e4)) | (e0 = op(e0, e3)) | (e0 = op(e0, e2)) | (e0 = op(e0, e1)) | (e0 = op(e0, e0))) & ((e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | (e1 = op(e4, e4)) | (e0 = op(e4, e4))) & ((e4 = op(e4, e3)) | (e3 = op(e4, e3)) | (e2 = op(e4, e3)) | (e1 = op(e4, e3)) | (e0 = op(e4, e3))) & ((e4 = op(e4, e2)) | (e3 = op(e4, e2)) | (e2 = op(e4, e2)) | (e1 = op(e4, e2)) | (e0 = op(e4, e2))) & ((e4 = op(e4, e1)) | (e3 = op(e4, e1)) | (e2 = op(e4, e1)) | (e1 = op(e4, e1)) | (e0 = op(e4, e1))) & ((e4 = op(e4, e0)) | (e3 = op(e4, e0)) | (e2 = op(e4, e0)) | (e1 = op(e4, e0)) | (e0 = op(e4, e0))) & ((e4 = op(e3, e4)) | (e3 = op(e3, e4)) | (e2 = op(e3, e4)) | (e1 = op(e3, e4)) | (e0 = op(e3, e4))) & ((e4 = op(e3, e3)) | (e3 = op(e3, e3)) | (e2 = op(e3, e3)) | (e1 = op(e3, e3)) | (e0 = op(e3, e3))) & ((e4 = op(e3, e2)) | (e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (e0 = op(e3, e2))) & ((e4 = op(e3, e1)) | (e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (e0 = op(e3, e1))) & ((e4 = op(e3, e0)) | (e3 = op(e3, e0)) | (e2 = op(e3, e0)) | (e1 = op(e3, e0)) | (e0 = op(e3, e0))) & ((e4 = op(e2, e4)) | (e3 = op(e2, e4)) | (e2 = op(e2, e4)) | (e1 = op(e2, e4)) | (e0 = op(e2, e4))) & ((e4 = op(e2, e3)) | (e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))) & ((e4 = op(e2, e2)) | (e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (e0 = op(e2, e2))) & ((e4 = op(e2, e1)) | (e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))) & ((e4 = op(e2, e0)) | (e3 = op(e2, e0)) | (e2 = op(e2, e0)) | (e1 = op(e2, e0)) | (e0 = op(e2, e0))) & ((e4 = op(e1, e4)) | (e3 = op(e1, e4)) | (e2 = op(e1, e4)) | (e1 = op(e1, e4)) | (e0 = op(e1, e4))) & ((e4 = op(e1, e3)) | (e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (e0 = op(e1, e3))) & ((e4 = op(e1, e2)) | (e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (e0 = op(e1, e2))) & ((e4 = op(e1, e1)) | (e3 = op(e1, e1)) | (e2 = op(e1, e1)) | (e1 = op(e1, e1)) | (e0 = op(e1, e1))) & ((e4 = op(e1, e0)) | (e3 = op(e1, e0)) | (e2 = op(e1, e0)) | (e1 = op(e1, e0)) | (e0 = op(e1, e0))) & ((e4 = op(e0, e4)) | (e3 = op(e0, e4)) | (e2 = op(e0, e4)) | (e1 = op(e0, e4)) | (e0 = op(e0, e4))) & ((e4 = op(e0, e3)) | (e3 = op(e0, e3)) | (e2 = op(e0, e3)) | (e1 = op(e0, e3)) | (e0 = op(e0, e3))) & ((e4 = op(e0, e2)) | (e3 = op(e0, e2)) | (e2 = op(e0, e2)) | (e1 = op(e0, e2)) | (e0 = op(e0, e2))) & ((e4 = op(e0, e1)) | (e3 = op(e0, e1)) | (e2 = op(e0, e1)) | (e1 = op(e0, e1)) | (e0 = op(e0, e1))) & ((e4 = op(e0, e0)) | (e3 = op(e0, e0)) | (e2 = op(e0, e0)) | (e1 = op(e0, e0)) | (e0 = op(e0, e0))) & ((~ (e4 = op(e4, e4)) & ~ (e3 = op(e4, e3)) & ~ (e2 = op(e4, e2)) & ~ (e1 = op(e4, e1)) & ~ (e0 = op(e4, e0))) | (~ (e4 = op(e3, e4)) & ~ (e3 = op(e3, e3)) & ~ (e2 = op(e3, e2)) & ~ (e1 = op(e3, e1)) & ~ (e0 = op(e3, e0))) | (~ (e4 = op(e2, e4)) & ~ (e3 = op(e2, e3)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e2, e1)) & ~ (e0 = op(e2, e0))) | (~ (e4 = op(e1, e4)) & ~ (e3 = op(e1, e3)) & ~ (e2 = op(e1, e2)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e1, e0))) | (~ (e4 = op(e0, e4)) & ~ (e3 = op(e0, e3)) & ~ (e2 = op(e0, e2)) & ~ (e1 = op(e0, e1)) & ~ (e0 = op(e0, e0))))), inference(negated_conjecture, [], [f3])).
fof(f3, plain, ~ ((e4 = op(op(e4, e4), op(e4, e4))) & (e4 = op(op(e3, e4), op(e4, e3))) & (e4 = op(op(e2, e4), op(e4, e2))) & (e4 = op(op(e1, e4), op(e4, e1))) & (e4 = op(op(e0, e4), op(e4, e0))) & (e3 = op(op(e4, e3), op(e3, e4))) & (e3 = op(op(e3, e3), op(e3, e3))) & (e3 = op(op(e2, e3), op(e3, e2))) & (e3 = op(op(e1, e3), op(e3, e1))) & (e3 = op(op(e0, e3), op(e3, e0))) & (e2 = op(op(e4, e2), op(e2, e4))) & (e2 = op(op(e3, e2), op(e2, e3))) & (e2 = op(op(e2, e2), op(e2, e2))) & (e2 = op(op(e1, e2), op(e2, e1))) & (e2 = op(op(e0, e2), op(e2, e0))) & (e1 = op(op(e4, e1), op(e1, e4))) & (e1 = op(op(e3, e1), op(e1, e3))) & (e1 = op(op(e2, e1), op(e1, e2))) & (e1 = op(op(e1, e1), op(e1, e1))) & (e1 = op(op(e0, e1), op(e1, e0))) & (e0 = op(op(e4, e0), op(e0, e4))) & (e0 = op(op(e3, e0), op(e0, e3))) & (e0 = op(op(e2, e0), op(e0, e2))) & (e0 = op(op(e1, e0), op(e0, e1))) & (e0 = op(op(e0, e0), op(e0, e0))) & ((e4 = op(e4, e4)) | (e4 = op(e3, e4)) | (e4 = op(e2, e4)) | (e4 = op(e1, e4)) | (e4 = op(e0, e4))) & ((e4 = op(e4, e4)) | (e4 = op(e4, e3)) | (e4 = op(e4, e2)) | (e4 = op(e4, e1)) | (e4 = op(e4, e0))) & ((e3 = op(e4, e4)) | (e3 = op(e3, e4)) | (e3 = op(e2, e4)) | (e3 = op(e1, e4)) | (e3 = op(e0, e4))) & ((e3 = op(e4, e4)) | (e3 = op(e4, e3)) | (e3 = op(e4, e2)) | (e3 = op(e4, e1)) | (e3 = op(e4, e0))) & ((e2 = op(e4, e4)) | (e2 = op(e3, e4)) | (e2 = op(e2, e4)) | (e2 = op(e1, e4)) | (e2 = op(e0, e4))) & ((e2 = op(e4, e4)) | (e2 = op(e4, e3)) | (e2 = op(e4, e2)) | (e2 = op(e4, e1)) | (e2 = op(e4, e0))) & ((e1 = op(e4, e4)) | (e1 = op(e3, e4)) | (e1 = op(e2, e4)) | (e1 = op(e1, e4)) | (e1 = op(e0, e4))) & ((e1 = op(e4, e4)) | (e1 = op(e4, e3)) | (e1 = op(e4, e2)) | (e1 = op(e4, e1)) | (e1 = op(e4, e0))) & ((e0 = op(e4, e4)) | (e0 = op(e3, e4)) | (e0 = op(e2, e4)) | (e0 = op(e1, e4)) | (e0 = op(e0, e4))) & ((e0 = op(e4, e4)) | (e0 = op(e4, e3)) | (e0 = op(e4, e2)) | (e0 = op(e4, e1)) | (e0 = op(e4, e0))) & ((e4 = op(e4, e3)) | (e4 = op(e3, e3)) | (e4 = op(e2, e3)) | (e4 = op(e1, e3)) | (e4 = op(e0, e3))) & ((e4 = op(e3, e4)) | (e4 = op(e3, e3)) | (e4 = op(e3, e2)) | (e4 = op(e3, e1)) | (e4 = op(e3, e0))) & ((e3 = op(e4, e3)) | (e3 = op(e3, e3)) | (e3 = op(e2, e3)) | (e3 = op(e1, e3)) | (e3 = op(e0, e3))) & ((e3 = op(e3, e4)) | (e3 = op(e3, e3)) | (e3 = op(e3, e2)) | (e3 = op(e3, e1)) | (e3 = op(e3, e0))) & ((e2 = op(e4, e3)) | (e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(e0, e3))) & ((e2 = op(e3, e4)) | (e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, e0))) & ((e1 = op(e4, e3)) | (e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(e0, e3))) & ((e1 = op(e3, e4)) | (e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, e0))) & ((e0 = op(e4, e3)) | (e0 = op(e3, e3)) | (e0 = op(e2, e3)) | (e0 = op(e1, e3)) | (e0 = op(e0, e3))) & ((e0 = op(e3, e4)) | (e0 = op(e3, e3)) | (e0 = op(e3, e2)) | (e0 = op(e3, e1)) | (e0 = op(e3, e0))) & ((e4 = op(e4, e2)) | (e4 = op(e3, e2)) | (e4 = op(e2, e2)) | (e4 = op(e1, e2)) | (e4 = op(e0, e2))) & ((e4 = op(e2, e4)) | (e4 = op(e2, e3)) | (e4 = op(e2, e2)) | (e4 = op(e2, e1)) | (e4 = op(e2, e0))) & ((e3 = op(e4, e2)) | (e3 = op(e3, e2)) | (e3 = op(e2, e2)) | (e3 = op(e1, e2)) | (e3 = op(e0, e2))) & ((e3 = op(e2, e4)) | (e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, e0))) & ((e2 = op(e4, e2)) | (e2 = op(e3, e2)) | (e2 = op(e2, e2)) | (e2 = op(e1, e2)) | (e2 = op(e0, e2))) & ((e2 = op(e2, e4)) | (e2 = op(e2, e3)) | (e2 = op(e2, e2)) | (e2 = op(e2, e1)) | (e2 = op(e2, e0))) & ((e1 = op(e4, e2)) | (e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(e0, e2))) & ((e1 = op(e2, e4)) | (e1 = op(e2, e3)) | (e1 = op(e2, e2)) | (e1 = op(e2, e1)) | (e1 = op(e2, e0))) & ((e0 = op(e4, e2)) | (e0 = op(e3, e2)) | (e0 = op(e2, e2)) | (e0 = op(e1, e2)) | (e0 = op(e0, e2))) & ((e0 = op(e2, e4)) | (e0 = op(e2, e3)) | (e0 = op(e2, e2)) | (e0 = op(e2, e1)) | (e0 = op(e2, e0))) & ((e4 = op(e4, e1)) | (e4 = op(e3, e1)) | (e4 = op(e2, e1)) | (e4 = op(e1, e1)) | (e4 = op(e0, e1))) & ((e4 = op(e1, e4)) | (e4 = op(e1, e3)) | (e4 = op(e1, e2)) | (e4 = op(e1, e1)) | (e4 = op(e1, e0))) & ((e3 = op(e4, e1)) | (e3 = op(e3, e1)) | (e3 = op(e2, e1)) | (e3 = op(e1, e1)) | (e3 = op(e0, e1))) & ((e3 = op(e1, e4)) | (e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, e0))) & ((e2 = op(e4, e1)) | (e2 = op(e3, e1)) | (e2 = op(e2, e1)) | (e2 = op(e1, e1)) | (e2 = op(e0, e1))) & ((e2 = op(e1, e4)) | (e2 = op(e1, e3)) | (e2 = op(e1, e2)) | (e2 = op(e1, e1)) | (e2 = op(e1, e0))) & ((e1 = op(e4, e1)) | (e1 = op(e3, e1)) | (e1 = op(e2, e1)) | (e1 = op(e1, e1)) | (e1 = op(e0, e1))) & ((e1 = op(e1, e4)) | (e1 = op(e1, e3)) | (e1 = op(e1, e2)) | (e1 = op(e1, e1)) | (e1 = op(e1, e0))) & ((e0 = op(e4, e1)) | (e0 = op(e3, e1)) | (e0 = op(e2, e1)) | (e0 = op(e1, e1)) | (e0 = op(e0, e1))) & ((e0 = op(e1, e4)) | (e0 = op(e1, e3)) | (e0 = op(e1, e2)) | (e0 = op(e1, e1)) | (e0 = op(e1, e0))) & ((e4 = op(e4, e0)) | (e4 = op(e3, e0)) | (e4 = op(e2, e0)) | (e4 = op(e1, e0)) | (e4 = op(e0, e0))) & ((e4 = op(e0, e4)) | (e4 = op(e0, e3)) | (e4 = op(e0, e2)) | (e4 = op(e0, e1)) | (e4 = op(e0, e0))) & ((e3 = op(e4, e0)) | (e3 = op(e3, e0)) | (e3 = op(e2, e0)) | (e3 = op(e1, e0)) | (e3 = op(e0, e0))) & ((e3 = op(e0, e4)) | (e3 = op(e0, e3)) | (e3 = op(e0, e2)) | (e3 = op(e0, e1)) | (e3 = op(e0, e0))) & ((e2 = op(e4, e0)) | (e2 = op(e3, e0)) | (e2 = op(e2, e0)) | (e2 = op(e1, e0)) | (e2 = op(e0, e0))) & ((e2 = op(e0, e4)) | (e2 = op(e0, e3)) | (e2 = op(e0, e2)) | (e2 = op(e0, e1)) | (e2 = op(e0, e0))) & ((e1 = op(e4, e0)) | (e1 = op(e3, e0)) | (e1 = op(e2, e0)) | (e1 = op(e1, e0)) | (e1 = op(e0, e0))) & ((e1 = op(e0, e4)) | (e1 = op(e0, e3)) | (e1 = op(e0, e2)) | (e1 = op(e0, e1)) | (e1 = op(e0, e0))) & ((e0 = op(e4, e0)) | (e0 = op(e3, e0)) | (e0 = op(e2, e0)) | (e0 = op(e1, e0)) | (e0 = op(e0, e0))) & ((e0 = op(e0, e4)) | (e0 = op(e0, e3)) | (e0 = op(e0, e2)) | (e0 = op(e0, e1)) | (e0 = op(e0, e0))) & ((e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | (e1 = op(e4, e4)) | (e0 = op(e4, e4))) & ((e4 = op(e4, e3)) | (e3 = op(e4, e3)) | (e2 = op(e4, e3)) | (e1 = op(e4, e3)) | (e0 = op(e4, e3))) & ((e4 = op(e4, e2)) | (e3 = op(e4, e2)) | (e2 = op(e4, e2)) | (e1 = op(e4, e2)) | (e0 = op(e4, e2))) & ((e4 = op(e4, e1)) | (e3 = op(e4, e1)) | (e2 = op(e4, e1)) | (e1 = op(e4, e1)) | (e0 = op(e4, e1))) & ((e4 = op(e4, e0)) | (e3 = op(e4, e0)) | (e2 = op(e4, e0)) | (e1 = op(e4, e0)) | (e0 = op(e4, e0))) & ((e4 = op(e3, e4)) | (e3 = op(e3, e4)) | (e2 = op(e3, e4)) | (e1 = op(e3, e4)) | (e0 = op(e3, e4))) & ((e4 = op(e3, e3)) | (e3 = op(e3, e3)) | (e2 = op(e3, e3)) | (e1 = op(e3, e3)) | (e0 = op(e3, e3))) & ((e4 = op(e3, e2)) | (e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (e0 = op(e3, e2))) & ((e4 = op(e3, e1)) | (e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (e0 = op(e3, e1))) & ((e4 = op(e3, e0)) | (e3 = op(e3, e0)) | (e2 = op(e3, e0)) | (e1 = op(e3, e0)) | (e0 = op(e3, e0))) & ((e4 = op(e2, e4)) | (e3 = op(e2, e4)) | (e2 = op(e2, e4)) | (e1 = op(e2, e4)) | (e0 = op(e2, e4))) & ((e4 = op(e2, e3)) | (e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))) & ((e4 = op(e2, e2)) | (e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (e0 = op(e2, e2))) & ((e4 = op(e2, e1)) | (e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))) & ((e4 = op(e2, e0)) | (e3 = op(e2, e0)) | (e2 = op(e2, e0)) | (e1 = op(e2, e0)) | (e0 = op(e2, e0))) & ((e4 = op(e1, e4)) | (e3 = op(e1, e4)) | (e2 = op(e1, e4)) | (e1 = op(e1, e4)) | (e0 = op(e1, e4))) & ((e4 = op(e1, e3)) | (e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (e0 = op(e1, e3))) & ((e4 = op(e1, e2)) | (e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (e0 = op(e1, e2))) & ((e4 = op(e1, e1)) | (e3 = op(e1, e1)) | (e2 = op(e1, e1)) | (e1 = op(e1, e1)) | (e0 = op(e1, e1))) & ((e4 = op(e1, e0)) | (e3 = op(e1, e0)) | (e2 = op(e1, e0)) | (e1 = op(e1, e0)) | (e0 = op(e1, e0))) & ((e4 = op(e0, e4)) | (e3 = op(e0, e4)) | (e2 = op(e0, e4)) | (e1 = op(e0, e4)) | (e0 = op(e0, e4))) & ((e4 = op(e0, e3)) | (e3 = op(e0, e3)) | (e2 = op(e0, e3)) | (e1 = op(e0, e3)) | (e0 = op(e0, e3))) & ((e4 = op(e0, e2)) | (e3 = op(e0, e2)) | (e2 = op(e0, e2)) | (e1 = op(e0, e2)) | (e0 = op(e0, e2))) & ((e4 = op(e0, e1)) | (e3 = op(e0, e1)) | (e2 = op(e0, e1)) | (e1 = op(e0, e1)) | (e0 = op(e0, e1))) & ((e4 = op(e0, e0)) | (e3 = op(e0, e0)) | (e2 = op(e0, e0)) | (e1 = op(e0, e0)) | (e0 = op(e0, e0))) & ((~ (e4 = op(e4, e4)) & ~ (e3 = op(e4, e3)) & ~ (e2 = op(e4, e2)) & ~ (e1 = op(e4, e1)) & ~ (e0 = op(e4, e0))) | (~ (e4 = op(e3, e4)) & ~ (e3 = op(e3, e3)) & ~ (e2 = op(e3, e2)) & ~ (e1 = op(e3, e1)) & ~ (e0 = op(e3, e0))) | (~ (e4 = op(e2, e4)) & ~ (e3 = op(e2, e3)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e2, e1)) & ~ (e0 = op(e2, e0))) | (~ (e4 = op(e1, e4)) & ~ (e3 = op(e1, e3)) & ~ (e2 = op(e1, e2)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e1, e0))) | (~ (e4 = op(e0, e4)) & ~ (e3 = op(e0, e3)) & ~ (e2 = op(e0, e2)) & ~ (e1 = op(e0, e1)) & ~ (e0 = op(e0, e0))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG187+1.p', co1)).