fof(f6730, plain, $false, inference(avatar_sat_refutation, [], [f401, f418, f435, f452, f469, f486, f503, f520, f537, f554, f571, f588, f605, f622, f639, f656, f657, f658, f659, f660, f661, f662, f663, f664, f665, f666, f667, f668, f669, f670, f671, f672, f673, f674, f675, f676, f677, f678, f679, f680, f681, f682, f683, f684, f685, f686, f687, f688, f702, f707, f712, f713, f714, f719, f720, f721, f722, f723, f724, f729, f730, f731, f733, f734, f743, f753, f758, f759, f760, f765, f767, f768, f770, f775, f778, f779, f780, f785, f786, f787, f788, f789, f790, f799, f804, f809, f814, f815, f816, f821, f822, f823, f824, f825, f826, f831, f832, f833, f834, f836, f841, f842, f844, f845, f846, f855, f860, f865, f870, f871, f872, f877, f878, f879, f880, f881, f882, f887, f888, f889, f890, f891, f892, f897, f898, f899, f900, f902, f911, f916, f917, f926, f931, f932, f941, f946, f947, f956, f961, f962, f971, f972, f973, f982, f987, f988, f997, f1002, f1003, f1012, f1017, f1018, f1027, f1032, f1033, f1042, f1043, f1053, f1058, f1059, f1068, f1073, f1074, f1083, f1088, f1089, f1098, f1103, f1104, f1113, f1114, f1115, f1116, f1117, f1118, f1119, f1121, f1126, f1127, f1128, f1137, f1146, f1155, f1160, f1165, f1170, f1179, f1188, f1197, f1198, f1203, f1204, f1213, f1218, f1227, f1228, f1229, f1230, f1235, f1240, f1245, f1246, f1247, f1248, f1263, f1293, f1296, f1303, f1306, f1337, f1340, f1355, f1379, f1382, f1385, f1394, f1397, f1404, f1407, f1415, f1416, f1419, f1422, f1425, f1430, f1433, f1437, f1450, f1453, f1484, f1489, f1526, f1538, f1552, f1554, f1562, f1565, f1580, f1588, f1589, f1613, f1621, f1629, f1634, f1672, f1683, f1689, f1693, f1726, f1739, f1802, f1805, f1813, f1816, f1829, f1839, f1852, f1878, f1903, f1906, f1913, f1932, f1940, f1949, f1969, f1986, f1994, f2007, f2020, f2024, f2026, f2058, f2075, f2078, f2084, f2087, f2099, f2112, f2117, f2123, f2125, f2131, f2132, f2149, f2150, f2160, f2178, f2182, f2185, f2194, f2228, f2241, f2266, f2275, f2278, f2283, f2284, f2297, f2312, f2318, f2320, f2337, f2366, f2376, f2383, f2399, f2424, f2446, f2464, f2477, f2484, f2491, f2500, f2527, f2537, f2540, f2543, f2563, f2570, f2586, f2596, f2622, f2651, f2654, f2718, f2728, f2747, f2750, f2758, f2762, f2781, f2791, f2798, f2821, f2824, f2831, f2846, f2847, f2852, f2874, f2924, f2926, f2949, f2986, f2996, f3000, f3012, f3024, f3035, f3038, f3054, f3070, f3100, f3101, f3102, f3108, f3133, f3156, f3158, f3181, f3202, f3228, f3231, f3238, f3254, f3261, f3264, f3278, f3285, f3288, f3314, f3333, f3366, f3407, f3411, f3420, f3422, f3430, f3433, f3462, f3488, f3496, f3538, f3541, f3549, f3555, f3583, f3589, f3593, f3600, f3609, f3611, f3616, f3674, f3698, f3703, f3722, f3743, f3758, f3767, f3771, f3778, f3780, f3790, f3791, f3800, f3850, f3861, f3907, f3927, f3929, f3943, f3944, f3950, f3951, f3957, f3976, f4039, f4087, f4104, f4124, f4145, f4162, f4164, f4165, f4172, f4174, f4189, f4193, f4222, f4227, f4241, f4244, f4245, f4249, f4271, f4274, f4276, f4278, f4287, f4296, f4297, f4305, f4306, f4356, f4366, f4377, f4380, f4384, f4409, f4416, f4430, f4431, f4438, f4495, f4508, f4533, f4541, f4554, f4557, f4576, f4583, f4596, f4598, f4604, f4618, f4634, f4639, f4657, f4659, f4691, f4707, f4754, f4757, f4760, f4795, f4806, f4824, f4828, f4835, f4840, f4875, f4876, f4877, f4892, f4893, f4904, f4918, f4942, f4958, f4963, f4974, f4978, f4994, f4996, f5015, f5016, f5018, f5036, f5037, f5039, f5053, f5109, f5117, f5132, f5133, f5140, f5151, f5158, f5169, f5206, f5216, f5217, f5231, f5238, f5264, f5266, f5281, f5283, f5309, f5315, f5330, f5345, f5362, f5371, f5396, f5398, f5400, f5419, f5430, f5440, f5451, f5452, f5461, f5475, f5478, f5479, f5486, f5498, f5516, f5517, f5528, f5530, f5540, f5561, f5563, f5565, f5567, f5569, f5582, f5588, f5590, f5600, f5608, f5619, f5626, f5654, f5681, f5682, f5684, f5689, f5698, f5706, f5708, f5721, f5741, f5742, f5751, f5764, f5798, f5808, f5826, f5828, f5830, f5834, f5842, f5843, f5844, f5854, f5860, f5871, f5883, f5891, f5893, f5898, f5900, f5907, f5921, f5927, f5931, f5940, f5941, f5951, f5954, f5962, f5964, f5982, f5983, f5989, f6014, f6025, f6027, f6031, f6035, f6038, f6047, f6048, f6049, f6060, f6061, f6071, f6079, f6088, f6089, f6091, f6097, f6105, f6115, f6144, f6148, f6150, f6159, f6160, f6161, f6170, f6172, f6181, f6195, f6196, f6213, f6227, f6240, f6249, f6274, f6280, f6282, f6289, f6290, f6292, f6299, f6309, f6328, f6330, f6334, f6347, f6351, f6354, f6362, f6363, f6364, f6373, f6385, f6391, f6392, f6393, f6400, f6409, f6410, f6421, f6429, f6447, f6455, f6457, f6459, f6461, f6469, f6470, f6471, f6483, f6484, f6492, f6503, f6504, f6513, f6514, f6516, f6536, f6542, f6543, f6544, f6552, f6560, f6567, f6580, f6585, f6597, f6599, f6601, f6603, f6618, f6619, f6635, f6636, f6637, f6645, f6646, f6657, f6673, f6674, f6675, f6682, f6692, f6695, f6704, f6705, f6714, f6724])).
fof(f6724, plain, (~ spl30_21 | ~ spl30_40 | ~ spl30_54 | spl30_161), inference(avatar_contradiction_clause, [], [f6723])).
fof(f6723, plain, ($false | (~ spl30_21 | ~ spl30_40 | ~ spl30_54 | spl30_161)), inference(subsumption_resolution, [], [f6722, f553])).
fof(f553, plain, ((e3 = op(e1, e2)) | ~ spl30_40), inference(avatar_component_clause, [], [f551])).
fof(f551, plain, (spl30_40 <=> (e3 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_40])])).
fof(f6722, plain, (~ (e3 = op(e1, e2)) | (~ spl30_21 | ~ spl30_54 | spl30_161)), inference(forward_demodulation, [], [f6721, f613])).
fof(f613, plain, ((e1 = op(e0, e2)) | ~ spl30_54), inference(avatar_component_clause, [], [f611])).
fof(f611, plain, (spl30_54 <=> (e1 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_54])])).
fof(f6721, plain, (~ (e3 = op(op(e0, e2), e2)) | (~ spl30_21 | spl30_161)), inference(forward_demodulation, [], [f1239, f473])).
fof(f473, plain, ((e0 = op(e2, e2)) | ~ spl30_21), inference(avatar_component_clause, [], [f471])).
fof(f471, plain, (spl30_21 <=> (e0 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_21])])).
fof(f1239, plain, (~ (e3 = op(op(op(e2, e2), e2), e2)) | spl30_161), inference(avatar_component_clause, [], [f1237])).
fof(f1237, plain, (spl30_161 <=> (e3 = op(op(op(e2, e2), e2), e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_161])])).
fof(f6714, plain, (~ spl30_5 | ~ spl30_9), inference(avatar_split_clause, [], [f6707, f420, f403])).
fof(f403, plain, (spl30_5 <=> (e0 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_5])])).
fof(f420, plain, (spl30_9 <=> (e0 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_9])])).
fof(f6707, plain, (~ (e0 = op(e3, e2)) | ~ spl30_9), inference(backward_demodulation, [], [f208, f422])).
fof(f422, plain, ((e0 = op(e3, e1)) | ~ spl30_9), inference(avatar_component_clause, [], [f420])).
fof(f208, plain, ~ (op(e3, e1) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f3, plain, (~ (op(e3, e2) = op(e3, e3)) & ~ (op(e3, e1) = op(e3, e3)) & ~ (op(e3, e1) = op(e3, e2)) & ~ (op(e3, e0) = op(e3, e3)) & ~ (op(e3, e0) = op(e3, e2)) & ~ (op(e3, e0) = op(e3, e1)) & ~ (op(e2, e2) = op(e2, e3)) & ~ (op(e2, e1) = op(e2, e3)) & ~ (op(e2, e1) = op(e2, e2)) & ~ (op(e2, e0) = op(e2, e3)) & ~ (op(e2, e0) = op(e2, e2)) & ~ (op(e2, e0) = op(e2, e1)) & ~ (op(e1, e2) = op(e1, e3)) & ~ (op(e1, e1) = op(e1, e3)) & ~ (op(e1, e1) = op(e1, e2)) & ~ (op(e1, e0) = op(e1, e3)) & ~ (op(e1, e0) = op(e1, e2)) & ~ (op(e1, e0) = op(e1, e1)) & ~ (op(e0, e2) = op(e0, e3)) & ~ (op(e0, e1) = op(e0, e3)) & ~ (op(e0, e1) = op(e0, e2)) & ~ (op(e0, e0) = op(e0, e3)) & ~ (op(e0, e0) = op(e0, e2)) & ~ (op(e0, e0) = op(e0, e1)) & ~ (op(e2, e3) = op(e3, e3)) & ~ (op(e1, e3) = op(e3, e3)) & ~ (op(e1, e3) = op(e2, e3)) & ~ (op(e0, e3) = op(e3, e3)) & ~ (op(e0, e3) = op(e2, e3)) & ~ (op(e0, e3) = op(e1, e3)) & ~ (op(e2, e2) = op(e3, e2)) & ~ (op(e1, e2) = op(e3, e2)) & ~ (op(e1, e2) = op(e2, e2)) & ~ (op(e0, e2) = op(e3, e2)) & ~ (op(e0, e2) = op(e2, e2)) & ~ (op(e0, e2) = op(e1, e2)) & ~ (op(e2, e1) = op(e3, e1)) & ~ (op(e1, e1) = op(e3, e1)) & ~ (op(e1, e1) = op(e2, e1)) & ~ (op(e0, e1) = op(e3, e1)) & ~ (op(e0, e1) = op(e2, e1)) & ~ (op(e0, e1) = op(e1, e1)) & ~ (op(e2, e0) = op(e3, e0)) & ~ (op(e1, e0) = op(e3, e0)) & ~ (op(e1, e0) = op(e2, e0)) & ~ (op(e0, e0) = op(e3, e0)) & ~ (op(e0, e0) = op(e2, e0)) & ~ (op(e0, e0) = op(e1, e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG164+1.p', ax3)).
fof(f6705, plain, (spl30_9 | ~ spl30_14 | ~ spl30_69), inference(avatar_split_clause, [], [f6701, f709, f441, f420])).
fof(f441, plain, (spl30_14 <=> (e1 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_14])])).
fof(f709, plain, (spl30_69 <=> (e0 = op(e3, op(e3, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl30_69])])).
fof(f6701, plain, ((e0 = op(e3, e1)) | (~ spl30_14 | ~ spl30_69)), inference(backward_demodulation, [], [f711, f443])).
fof(f443, plain, ((e1 = op(e3, e0)) | ~ spl30_14), inference(avatar_component_clause, [], [f441])).
fof(f711, plain, ((e0 = op(e3, op(e3, e0))) | ~ spl30_69), inference(avatar_component_clause, [], [f709])).
fof(f6704, plain, (~ spl30_10 | ~ spl30_14), inference(avatar_split_clause, [], [f6697, f441, f424])).
fof(f424, plain, (spl30_10 <=> (e1 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_10])])).
fof(f6697, plain, (~ (e1 = op(e3, e1)) | ~ spl30_14), inference(backward_demodulation, [], [f205, f443])).
fof(f205, plain, ~ (op(e3, e0) = op(e3, e1)), inference(cnf_transformation, [], [f3])).
fof(f6695, plain, (~ spl30_21 | ~ spl30_54 | spl30_140), inference(avatar_contradiction_clause, [], [f6694])).
fof(f6694, plain, ($false | (~ spl30_21 | ~ spl30_54 | spl30_140)), inference(subsumption_resolution, [], [f6689, f613])).
fof(f6689, plain, (~ (e1 = op(e0, e2)) | (~ spl30_21 | spl30_140)), inference(backward_demodulation, [], [f1136, f473])).
fof(f1136, plain, (~ (e1 = op(op(e2, e2), e2)) | spl30_140), inference(avatar_component_clause, [], [f1134])).
fof(f1134, plain, (spl30_140 <=> (e1 = op(op(e2, e2), e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_140])])).
fof(f6692, plain, (~ spl30_5 | ~ spl30_21), inference(avatar_split_clause, [], [f6683, f471, f403])).
fof(f6683, plain, (~ (e0 = op(e3, e2)) | ~ spl30_21), inference(backward_demodulation, [], [f180, f473])).
fof(f180, plain, ~ (op(e2, e2) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f6682, plain, (~ spl30_23 | ~ spl30_27), inference(avatar_split_clause, [], [f6677, f496, f479])).
fof(f479, plain, (spl30_23 <=> (e2 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_23])])).
fof(f496, plain, (spl30_27 <=> (e2 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_27])])).
fof(f6677, plain, (~ (e2 = op(e2, e2)) | ~ spl30_27), inference(backward_demodulation, [], [f202, f498])).
fof(f498, plain, ((e2 = op(e2, e1)) | ~ spl30_27), inference(avatar_component_clause, [], [f496])).
fof(f202, plain, ~ (op(e2, e1) = op(e2, e2)), inference(cnf_transformation, [], [f3])).
fof(f6675, plain, (~ spl30_5 | ~ spl30_32 | spl30_115), inference(avatar_split_clause, [], [f6672, f999, f517, f403])).
fof(f517, plain, (spl30_32 <=> (e3 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_32])])).
fof(f999, plain, (spl30_115 <=> (e0 = op(op(e2, e0), e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_115])])).
fof(f6672, plain, (~ (e0 = op(e3, e2)) | (~ spl30_32 | spl30_115)), inference(backward_demodulation, [], [f1000, f519])).
fof(f519, plain, ((e3 = op(e2, e0)) | ~ spl30_32), inference(avatar_component_clause, [], [f517])).
fof(f1000, plain, (~ (e0 = op(op(e2, e0), e2)) | spl30_115), inference(avatar_component_clause, [], [f999])).
fof(f6674, plain, (~ spl30_15 | ~ spl30_32 | spl30_114), inference(avatar_split_clause, [], [f6671, f994, f517, f445])).
fof(f445, plain, (spl30_15 <=> (e2 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_15])])).
fof(f994, plain, (spl30_114 <=> (e2 = op(op(e2, e0), e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_114])])).
fof(f6671, plain, (~ (e2 = op(e3, e0)) | (~ spl30_32 | spl30_114)), inference(backward_demodulation, [], [f996, f519])).
fof(f996, plain, (~ (e2 = op(op(e2, e0), e0)) | spl30_114), inference(avatar_component_clause, [], [f994])).
fof(f6673, plain, (~ spl30_28 | ~ spl30_32), inference(avatar_split_clause, [], [f6668, f517, f500])).
fof(f500, plain, (spl30_28 <=> (e3 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_28])])).
fof(f6668, plain, (~ (e3 = op(e2, e1)) | ~ spl30_32), inference(backward_demodulation, [], [f199, f519])).
fof(f199, plain, ~ (op(e2, e0) = op(e2, e1)), inference(cnf_transformation, [], [f3])).
fof(f6657, plain, (~ spl30_10 | ~ spl30_42), inference(avatar_split_clause, [], [f6648, f560, f424])).
fof(f560, plain, (spl30_42 <=> (e1 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_42])])).
fof(f6648, plain, (~ (e1 = op(e3, e1)) | ~ spl30_42), inference(backward_demodulation, [], [f173, f562])).
fof(f562, plain, ((e1 = op(e1, e1)) | ~ spl30_42), inference(avatar_component_clause, [], [f560])).
fof(f173, plain, ~ (op(e1, e1) = op(e3, e1)), inference(cnf_transformation, [], [f3])).
fof(f6646, plain, (~ spl30_43 | ~ spl30_47), inference(avatar_split_clause, [], [f6640, f581, f564])).
fof(f564, plain, (spl30_43 <=> (e2 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_43])])).
fof(f581, plain, (spl30_47 <=> (e2 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_47])])).
fof(f6640, plain, (~ (e2 = op(e1, e1)) | ~ spl30_47), inference(backward_demodulation, [], [f193, f583])).
fof(f583, plain, ((e2 = op(e1, e0)) | ~ spl30_47), inference(avatar_component_clause, [], [f581])).
fof(f193, plain, ~ (op(e1, e0) = op(e1, e1)), inference(cnf_transformation, [], [f3])).
fof(f6645, plain, (~ spl30_15 | ~ spl30_47), inference(avatar_split_clause, [], [f6639, f581, f445])).
fof(f6639, plain, (~ (e2 = op(e3, e0)) | ~ spl30_47), inference(backward_demodulation, [], [f167, f583])).
fof(f167, plain, ~ (op(e1, e0) = op(e3, e0)), inference(cnf_transformation, [], [f3])).
fof(f6637, plain, (spl30_14 | ~ spl30_60 | ~ spl30_135), inference(avatar_split_clause, [], [f6634, f1100, f636, f441])).
fof(f636, plain, (spl30_60 <=> (e3 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_60])])).
fof(f1100, plain, (spl30_135 <=> (e1 = op(op(e0, e1), e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_135])])).
fof(f6634, plain, ((e1 = op(e3, e0)) | (~ spl30_60 | ~ spl30_135)), inference(backward_demodulation, [], [f1102, f638])).
fof(f638, plain, ((e3 = op(e0, e1)) | ~ spl30_60), inference(avatar_component_clause, [], [f636])).
fof(f1102, plain, ((e1 = op(op(e0, e1), e0)) | ~ spl30_135), inference(avatar_component_clause, [], [f1100])).
fof(f6636, plain, (spl30_9 | ~ spl30_60 | ~ spl30_134), inference(avatar_split_clause, [], [f6633, f1095, f636, f420])).
fof(f1095, plain, (spl30_134 <=> (e0 = op(op(e0, e1), e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_134])])).
fof(f6633, plain, ((e0 = op(e3, e1)) | (~ spl30_60 | ~ spl30_134)), inference(backward_demodulation, [], [f1096, f638])).
fof(f1096, plain, ((e0 = op(op(e0, e1), e1)) | ~ spl30_134), inference(avatar_component_clause, [], [f1095])).
fof(f6635, plain, (~ spl30_28 | ~ spl30_60), inference(avatar_split_clause, [], [f6630, f636, f500])).
fof(f6630, plain, (~ (e3 = op(e2, e1)) | ~ spl30_60), inference(backward_demodulation, [], [f170, f638])).
fof(f170, plain, ~ (op(e0, e1) = op(e2, e1)), inference(cnf_transformation, [], [f3])).
fof(f6619, plain, (~ spl30_57 | ~ spl30_61), inference(avatar_split_clause, [], [f6607, f641, f624])).
fof(f624, plain, (spl30_57 <=> (e0 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_57])])).
fof(f641, plain, (spl30_61 <=> (e0 = op(e0, e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_61])])).
fof(f6607, plain, (~ (e0 = op(e0, e1)) | ~ spl30_61), inference(backward_demodulation, [], [f187, f643])).
fof(f643, plain, ((e0 = op(e0, e0)) | ~ spl30_61), inference(avatar_component_clause, [], [f641])).
fof(f187, plain, ~ (op(e0, e0) = op(e0, e1)), inference(cnf_transformation, [], [f3])).
fof(f6618, plain, (~ spl30_29 | ~ spl30_61), inference(avatar_split_clause, [], [f6605, f641, f505])).
fof(f505, plain, (spl30_29 <=> (e0 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_29])])).
fof(f6605, plain, (~ (e0 = op(e2, e0)) | ~ spl30_61), inference(backward_demodulation, [], [f164, f643])).
fof(f164, plain, ~ (op(e0, e0) = op(e2, e0)), inference(cnf_transformation, [], [f3])).
fof(f6603, plain, (~ spl30_28 | ~ spl30_18 | spl30_73), inference(avatar_split_clause, [], [f6602, f740, f458, f500])).
fof(f458, plain, (spl30_18 <=> (e1 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_18])])).
fof(f740, plain, (spl30_73 <=> (e3 = op(e2, op(e2, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl30_73])])).
fof(f6602, plain, (~ (e3 = op(e2, e1)) | (~ spl30_18 | spl30_73)), inference(forward_demodulation, [], [f741, f460])).
fof(f460, plain, ((e1 = op(e2, e3)) | ~ spl30_18), inference(avatar_component_clause, [], [f458])).
fof(f741, plain, (~ (e3 = op(e2, op(e2, e3))) | spl30_73), inference(avatar_component_clause, [], [f740])).
fof(f6601, plain, (spl30_60 | ~ spl30_33 | ~ spl30_118), inference(avatar_split_clause, [], [f6600, f1014, f522, f636])).
fof(f522, plain, (spl30_33 <=> (e0 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_33])])).
fof(f1014, plain, (spl30_118 <=> (e3 = op(op(e1, e3), e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_118])])).
fof(f6600, plain, ((e3 = op(e0, e1)) | (~ spl30_33 | ~ spl30_118)), inference(forward_demodulation, [], [f1016, f524])).
fof(f524, plain, ((e0 = op(e1, e3)) | ~ spl30_33), inference(avatar_component_clause, [], [f522])).
fof(f1016, plain, ((e3 = op(op(e1, e3), e1)) | ~ spl30_118), inference(avatar_component_clause, [], [f1014])).
fof(f6599, plain, (spl30_32 | ~ spl30_51 | ~ spl30_129), inference(avatar_split_clause, [], [f6598, f1070, f598, f517])).
fof(f598, plain, (spl30_51 <=> (e2 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_51])])).
fof(f1070, plain, (spl30_129 <=> (e3 = op(op(e0, e3), e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_129])])).
fof(f6598, plain, ((e3 = op(e2, e0)) | (~ spl30_51 | ~ spl30_129)), inference(forward_demodulation, [], [f1072, f600])).
fof(f600, plain, ((e2 = op(e0, e3)) | ~ spl30_51), inference(avatar_component_clause, [], [f598])).
fof(f1072, plain, ((e3 = op(op(e0, e3), e0)) | ~ spl30_129), inference(avatar_component_clause, [], [f1070])).
fof(f6597, plain, (spl30_47 | ~ spl30_54 | ~ spl30_132), inference(avatar_split_clause, [], [f6596, f1085, f611, f581])).
fof(f1085, plain, (spl30_132 <=> (e2 = op(op(e0, e2), e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_132])])).
fof(f6596, plain, ((e2 = op(e1, e0)) | (~ spl30_54 | ~ spl30_132)), inference(forward_demodulation, [], [f1087, f613])).
fof(f1087, plain, ((e2 = op(op(e0, e2), e0)) | ~ spl30_132), inference(avatar_component_clause, [], [f1085])).
fof(f6585, plain, (~ spl30_48 | ~ spl30_33 | spl30_81), inference(avatar_split_clause, [], [f6522, f796, f522, f585])).
fof(f585, plain, (spl30_48 <=> (e3 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_48])])).
fof(f796, plain, (spl30_81 <=> (e3 = op(e1, op(e1, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl30_81])])).
fof(f6522, plain, (~ (e3 = op(e1, e0)) | (~ spl30_33 | spl30_81)), inference(backward_demodulation, [], [f797, f524])).
fof(f797, plain, (~ (e3 = op(e1, op(e1, e3))) | spl30_81), inference(avatar_component_clause, [], [f796])).
fof(f6580, plain, (~ spl30_24 | ~ spl30_40), inference(avatar_split_clause, [], [f6579, f551, f483])).
fof(f483, plain, (spl30_24 <=> (e3 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_24])])).
fof(f6579, plain, (~ (e3 = op(e2, e2)) | ~ spl30_40), inference(forward_demodulation, [], [f178, f553])).
fof(f178, plain, ~ (op(e1, e2) = op(e2, e2)), inference(cnf_transformation, [], [f3])).
fof(f6567, plain, (~ spl30_2 | ~ spl30_10), inference(avatar_split_clause, [], [f6562, f424, f390])).
fof(f390, plain, (spl30_2 <=> (e1 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_2])])).
fof(f6562, plain, (~ (e1 = op(e3, e3)) | ~ spl30_10), inference(backward_demodulation, [], [f209, f426])).
fof(f426, plain, ((e1 = op(e3, e1)) | ~ spl30_10), inference(avatar_component_clause, [], [f424])).
fof(f209, plain, ~ (op(e3, e1) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f6560, plain, (~ spl30_11 | ~ spl30_15), inference(avatar_split_clause, [], [f6554, f445, f428])).
fof(f428, plain, (spl30_11 <=> (e2 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_11])])).
fof(f6554, plain, (~ (e2 = op(e3, e1)) | ~ spl30_15), inference(backward_demodulation, [], [f205, f447])).
fof(f447, plain, ((e2 = op(e3, e0)) | ~ spl30_15), inference(avatar_component_clause, [], [f445])).
fof(f6552, plain, (~ spl30_2 | ~ spl30_18), inference(avatar_split_clause, [], [f6546, f458, f390])).
fof(f6546, plain, (~ (e1 = op(e3, e3)) | ~ spl30_18), inference(backward_demodulation, [], [f186, f460])).
fof(f186, plain, ~ (op(e2, e3) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f6544, plain, (~ spl30_11 | ~ spl30_28 | spl30_111), inference(avatar_split_clause, [], [f6540, f979, f500, f428])).
fof(f979, plain, (spl30_111 <=> (e2 = op(op(e2, e1), e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_111])])).
fof(f6540, plain, (~ (e2 = op(e3, e1)) | (~ spl30_28 | spl30_111)), inference(backward_demodulation, [], [f981, f502])).
fof(f502, plain, ((e3 = op(e2, e1)) | ~ spl30_28), inference(avatar_component_clause, [], [f500])).
fof(f981, plain, (~ (e2 = op(op(e2, e1), e1)) | spl30_111), inference(avatar_component_clause, [], [f979])).
fof(f6543, plain, (spl30_18 | ~ spl30_28 | ~ spl30_75), inference(avatar_split_clause, [], [f6539, f750, f500, f458])).
fof(f750, plain, (spl30_75 <=> (e1 = op(e2, op(e2, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl30_75])])).
fof(f6539, plain, ((e1 = op(e2, e3)) | (~ spl30_28 | ~ spl30_75)), inference(backward_demodulation, [], [f752, f502])).
fof(f752, plain, ((e1 = op(e2, op(e2, e1))) | ~ spl30_75), inference(avatar_component_clause, [], [f750])).
fof(f6542, plain, (~ spl30_20 | ~ spl30_28), inference(avatar_split_clause, [], [f6538, f500, f466])).
fof(f466, plain, (spl30_20 <=> (e3 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_20])])).
fof(f6538, plain, (~ (e3 = op(e2, e3)) | ~ spl30_28), inference(backward_demodulation, [], [f203, f502])).
fof(f203, plain, ~ (op(e2, e1) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f6536, plain, (~ spl30_25 | ~ spl30_29), inference(avatar_split_clause, [], [f6528, f505, f488])).
fof(f488, plain, (spl30_25 <=> (e0 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_25])])).
fof(f6528, plain, (~ (e0 = op(e2, e1)) | ~ spl30_29), inference(backward_demodulation, [], [f199, f507])).
fof(f507, plain, ((e0 = op(e2, e0)) | ~ spl30_29), inference(avatar_component_clause, [], [f505])).
fof(f6516, plain, (~ spl30_25 | ~ spl30_43 | spl30_156), inference(avatar_split_clause, [], [f6512, f1210, f564, f488])).
fof(f1210, plain, (spl30_156 <=> (e0 = op(op(e1, e1), e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_156])])).
fof(f6512, plain, (~ (e0 = op(e2, e1)) | (~ spl30_43 | spl30_156)), inference(backward_demodulation, [], [f1212, f566])).
fof(f566, plain, ((e2 = op(e1, e1)) | ~ spl30_43), inference(avatar_component_clause, [], [f564])).
fof(f1212, plain, (~ (e0 = op(op(e1, e1), e1)) | spl30_156), inference(avatar_component_clause, [], [f1210])).
fof(f6514, plain, (~ spl30_35 | ~ spl30_43), inference(avatar_split_clause, [], [f6507, f564, f530])).
fof(f530, plain, (spl30_35 <=> (e2 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_35])])).
fof(f6507, plain, (~ (e2 = op(e1, e3)) | ~ spl30_43), inference(backward_demodulation, [], [f197, f566])).
fof(f197, plain, ~ (op(e1, e1) = op(e1, e3)), inference(cnf_transformation, [], [f3])).
fof(f6513, plain, (~ spl30_11 | ~ spl30_43), inference(avatar_split_clause, [], [f6506, f564, f428])).
fof(f6506, plain, (~ (e2 = op(e3, e1)) | ~ spl30_43), inference(backward_demodulation, [], [f173, f566])).
fof(f6504, plain, (~ spl30_42 | ~ spl30_46), inference(avatar_split_clause, [], [f6496, f577, f560])).
fof(f577, plain, (spl30_46 <=> (e1 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_46])])).
fof(f6496, plain, (~ (e1 = op(e1, e1)) | ~ spl30_46), inference(backward_demodulation, [], [f193, f579])).
fof(f579, plain, ((e1 = op(e1, e0)) | ~ spl30_46), inference(avatar_component_clause, [], [f577])).
fof(f6503, plain, (~ spl30_30 | ~ spl30_46), inference(avatar_split_clause, [], [f6494, f577, f509])).
fof(f509, plain, (spl30_30 <=> (e1 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_30])])).
fof(f6494, plain, (~ (e1 = op(e2, e0)) | ~ spl30_46), inference(backward_demodulation, [], [f166, f579])).
fof(f166, plain, ~ (op(e1, e0) = op(e2, e0)), inference(cnf_transformation, [], [f3])).
fof(f6492, plain, (~ spl30_35 | ~ spl30_51), inference(avatar_split_clause, [], [f6485, f598, f530])).
fof(f6485, plain, (~ (e2 = op(e1, e3)) | ~ spl30_51), inference(backward_demodulation, [], [f181, f600])).
fof(f181, plain, ~ (op(e0, e3) = op(e1, e3)), inference(cnf_transformation, [], [f3])).
fof(f6484, plain, (~ spl30_49 | ~ spl30_57), inference(avatar_split_clause, [], [f6476, f624, f590])).
fof(f590, plain, (spl30_49 <=> (e0 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_49])])).
fof(f6476, plain, (~ (e0 = op(e0, e3)) | ~ spl30_57), inference(backward_demodulation, [], [f191, f626])).
fof(f626, plain, ((e0 = op(e0, e1)) | ~ spl30_57), inference(avatar_component_clause, [], [f624])).
fof(f191, plain, ~ (op(e0, e1) = op(e0, e3)), inference(cnf_transformation, [], [f3])).
fof(f6483, plain, (~ spl30_25 | ~ spl30_57), inference(avatar_split_clause, [], [f6474, f624, f488])).
fof(f6474, plain, (~ (e0 = op(e2, e1)) | ~ spl30_57), inference(backward_demodulation, [], [f170, f626])).
fof(f6471, plain, (~ spl30_49 | ~ spl30_64 | spl30_92), inference(avatar_split_clause, [], [f6467, f867, f653, f590])).
fof(f653, plain, (spl30_64 <=> (op(e0, e0) = e3)), introduced(avatar_definition, [new_symbols(naming, [spl30_64])])).
fof(f867, plain, (spl30_92 <=> (e0 = op(e0, op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl30_92])])).
fof(f6467, plain, (~ (e0 = op(e0, e3)) | (~ spl30_64 | spl30_92)), inference(backward_demodulation, [], [f868, f655])).
fof(f655, plain, ((op(e0, e0) = e3) | ~ spl30_64), inference(avatar_component_clause, [], [f653])).
fof(f868, plain, (~ (e0 = op(e0, op(e0, e0))) | spl30_92), inference(avatar_component_clause, [], [f867])).
fof(f6470, plain, (~ spl30_60 | ~ spl30_64), inference(avatar_split_clause, [], [f6465, f653, f636])).
fof(f6465, plain, (~ (e3 = op(e0, e1)) | ~ spl30_64), inference(backward_demodulation, [], [f187, f655])).
fof(f6469, plain, (~ spl30_16 | ~ spl30_64), inference(avatar_split_clause, [], [f6464, f653, f449])).
fof(f449, plain, (spl30_16 <=> (e3 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_16])])).
fof(f6464, plain, (~ (e3 = op(e3, e0)) | ~ spl30_64), inference(backward_demodulation, [], [f165, f655])).
fof(f165, plain, ~ (op(e0, e0) = op(e3, e0)), inference(cnf_transformation, [], [f3])).
fof(f6461, plain, (~ spl30_15 | ~ spl30_5 | spl30_67), inference(avatar_split_clause, [], [f6460, f699, f403, f445])).
fof(f699, plain, (spl30_67 <=> (e2 = op(e3, op(e3, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl30_67])])).
fof(f6460, plain, (~ (e2 = op(e3, e0)) | (~ spl30_5 | spl30_67)), inference(forward_demodulation, [], [f700, f405])).
fof(f405, plain, ((e0 = op(e3, e2)) | ~ spl30_5), inference(avatar_component_clause, [], [f403])).
fof(f700, plain, (~ (e2 = op(e3, op(e3, e2))) | spl30_67), inference(avatar_component_clause, [], [f699])).
fof(f6459, plain, (~ spl30_35 | ~ spl30_40 | spl30_82), inference(avatar_split_clause, [], [f6458, f801, f551, f530])).
fof(f801, plain, (spl30_82 <=> (e2 = op(e1, op(e1, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl30_82])])).
fof(f6458, plain, (~ (e2 = op(e1, e3)) | (~ spl30_40 | spl30_82)), inference(forward_demodulation, [], [f802, f553])).
fof(f802, plain, (~ (e2 = op(e1, op(e1, e2))) | spl30_82), inference(avatar_component_clause, [], [f801])).
fof(f6457, plain, (spl30_51 | ~ spl30_5 | ~ spl30_98), inference(avatar_split_clause, [], [f6456, f913, f403, f598])).
fof(f913, plain, (spl30_98 <=> (e2 = op(op(e3, e2), e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_98])])).
fof(f6456, plain, ((e2 = op(e0, e3)) | (~ spl30_5 | ~ spl30_98)), inference(forward_demodulation, [], [f915, f405])).
fof(f915, plain, ((e2 = op(op(e3, e2), e3)) | ~ spl30_98), inference(avatar_component_clause, [], [f913])).
fof(f6455, plain, (~ spl30_11 | ~ spl30_40 | spl30_121), inference(avatar_split_clause, [], [f6454, f1029, f551, f428])).
fof(f1029, plain, (spl30_121 <=> (e2 = op(op(e1, e2), e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_121])])).
fof(f6454, plain, (~ (e2 = op(e3, e1)) | (~ spl30_40 | spl30_121)), inference(forward_demodulation, [], [f1030, f553])).
fof(f1030, plain, (~ (e2 = op(op(e1, e2), e1)) | spl30_121), inference(avatar_component_clause, [], [f1029])).
fof(f6447, plain, (~ spl30_30 | spl30_152 | ~ spl30_153), inference(avatar_split_clause, [], [f6446, f1194, f1190, f509])).
fof(f1190, plain, (spl30_152 <=> (e1 = op(op(op(e0, e0), e0), e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_152])])).
fof(f1194, plain, (spl30_153 <=> (e2 = op(op(e0, e0), e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_153])])).
fof(f6446, plain, (~ (e1 = op(e2, e0)) | (spl30_152 | ~ spl30_153)), inference(forward_demodulation, [], [f1192, f1195])).
fof(f1195, plain, ((e2 = op(op(e0, e0), e0)) | ~ spl30_153), inference(avatar_component_clause, [], [f1194])).
fof(f1192, plain, (~ (e1 = op(op(op(e0, e0), e0), e0)) | spl30_152), inference(avatar_component_clause, [], [f1190])).
fof(f6429, plain, (~ spl30_1 | ~ spl30_5), inference(avatar_split_clause, [], [f6423, f403, f386])).
fof(f386, plain, (spl30_1 <=> (e0 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_1])])).
fof(f6423, plain, (~ (e0 = op(e3, e3)) | ~ spl30_5), inference(backward_demodulation, [], [f210, f405])).
fof(f210, plain, ~ (op(e3, e2) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f6421, plain, (~ spl30_8 | ~ spl30_16), inference(avatar_split_clause, [], [f6417, f449, f415])).
fof(f415, plain, (spl30_8 <=> (e3 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_8])])).
fof(f6417, plain, (~ (e3 = op(e3, e2)) | ~ spl30_16), inference(backward_demodulation, [], [f206, f451])).
fof(f451, plain, ((e3 = op(e3, e0)) | ~ spl30_16), inference(avatar_component_clause, [], [f449])).
fof(f206, plain, ~ (op(e3, e0) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f6410, plain, (~ spl30_18 | ~ spl30_30), inference(avatar_split_clause, [], [f6403, f509, f458])).
fof(f6403, plain, (~ (e1 = op(e2, e3)) | ~ spl30_30), inference(backward_demodulation, [], [f201, f511])).
fof(f511, plain, ((e1 = op(e2, e0)) | ~ spl30_30), inference(avatar_component_clause, [], [f509])).
fof(f201, plain, ~ (op(e2, e0) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f6409, plain, (~ spl30_14 | ~ spl30_30), inference(avatar_split_clause, [], [f6402, f509, f441])).
fof(f6402, plain, (~ (e1 = op(e3, e0)) | ~ spl30_30), inference(backward_demodulation, [], [f168, f511])).
fof(f168, plain, ~ (op(e2, e0) = op(e3, e0)), inference(cnf_transformation, [], [f3])).
fof(f6400, plain, (~ spl30_18 | ~ spl30_35 | spl30_117), inference(avatar_split_clause, [], [f6398, f1009, f530, f458])).
fof(f1009, plain, (spl30_117 <=> (e1 = op(op(e1, e3), e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_117])])).
fof(f6398, plain, (~ (e1 = op(e2, e3)) | (~ spl30_35 | spl30_117)), inference(backward_demodulation, [], [f1011, f532])).
fof(f532, plain, ((e2 = op(e1, e3)) | ~ spl30_35), inference(avatar_component_clause, [], [f530])).
fof(f1011, plain, (~ (e1 = op(op(e1, e3), e3)) | spl30_117), inference(avatar_component_clause, [], [f1009])).
fof(f6393, plain, (spl30_35 | ~ spl30_40 | ~ spl30_82), inference(avatar_split_clause, [], [f6388, f801, f551, f530])).
fof(f6388, plain, ((e2 = op(e1, e3)) | (~ spl30_40 | ~ spl30_82)), inference(backward_demodulation, [], [f803, f553])).
fof(f803, plain, ((e2 = op(e1, op(e1, e2))) | ~ spl30_82), inference(avatar_component_clause, [], [f801])).
fof(f6392, plain, (~ spl30_36 | ~ spl30_40), inference(avatar_split_clause, [], [f6387, f551, f534])).
fof(f534, plain, (spl30_36 <=> (e3 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_36])])).
fof(f6387, plain, (~ (e3 = op(e1, e3)) | ~ spl30_40), inference(backward_demodulation, [], [f198, f553])).
fof(f198, plain, ~ (op(e1, e2) = op(e1, e3)), inference(cnf_transformation, [], [f3])).
fof(f6391, plain, (~ spl30_8 | ~ spl30_40), inference(avatar_split_clause, [], [f6386, f551, f415])).
fof(f6386, plain, (~ (e3 = op(e3, e2)) | ~ spl30_40), inference(backward_demodulation, [], [f179, f553])).
fof(f179, plain, ~ (op(e1, e2) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f6385, plain, (~ spl30_37 | ~ spl30_45), inference(avatar_split_clause, [], [f6377, f573, f539])).
fof(f539, plain, (spl30_37 <=> (e0 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_37])])).
fof(f573, plain, (spl30_45 <=> (e0 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_45])])).
fof(f6377, plain, (~ (e0 = op(e1, e2)) | ~ spl30_45), inference(backward_demodulation, [], [f194, f575])).
fof(f575, plain, ((e0 = op(e1, e0)) | ~ spl30_45), inference(avatar_component_clause, [], [f573])).
fof(f194, plain, ~ (op(e1, e0) = op(e1, e2)), inference(cnf_transformation, [], [f3])).
fof(f6373, plain, (~ spl30_1 | ~ spl30_49), inference(avatar_split_clause, [], [f6367, f590, f386])).
fof(f6367, plain, (~ (e0 = op(e3, e3)) | ~ spl30_49), inference(backward_demodulation, [], [f183, f592])).
fof(f592, plain, ((e0 = op(e0, e3)) | ~ spl30_49), inference(avatar_component_clause, [], [f590])).
fof(f183, plain, ~ (op(e0, e3) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f6364, plain, (spl30_30 | ~ spl30_63 | ~ spl30_159), inference(avatar_split_clause, [], [f6361, f1224, f649, f509])).
fof(f649, plain, (spl30_63 <=> (op(e0, e0) = e2)), introduced(avatar_definition, [new_symbols(naming, [spl30_63])])).
fof(f1224, plain, (spl30_159 <=> (e1 = op(op(e0, e0), e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_159])])).
fof(f6361, plain, ((e1 = op(e2, e0)) | (~ spl30_63 | ~ spl30_159)), inference(backward_demodulation, [], [f1225, f651])).
fof(f651, plain, ((op(e0, e0) = e2) | ~ spl30_63), inference(avatar_component_clause, [], [f649])).
fof(f1225, plain, ((e1 = op(op(e0, e0), e0)) | ~ spl30_159), inference(avatar_component_clause, [], [f1224])).
fof(f6363, plain, (~ spl30_51 | ~ spl30_63), inference(avatar_split_clause, [], [f6358, f649, f598])).
fof(f6358, plain, (~ (e2 = op(e0, e3)) | ~ spl30_63), inference(backward_demodulation, [], [f189, f651])).
fof(f189, plain, ~ (op(e0, e0) = op(e0, e3)), inference(cnf_transformation, [], [f3])).
fof(f6362, plain, (~ spl30_47 | ~ spl30_63), inference(avatar_split_clause, [], [f6355, f649, f581])).
fof(f6355, plain, (~ (e2 = op(e1, e0)) | ~ spl30_63), inference(backward_demodulation, [], [f163, f651])).
fof(f163, plain, ~ (op(e0, e0) = op(e1, e0)), inference(cnf_transformation, [], [f3])).
fof(f6354, plain, (~ spl30_30 | ~ spl30_25 | spl30_75), inference(avatar_split_clause, [], [f6353, f750, f488, f509])).
fof(f6353, plain, (~ (e1 = op(e2, e0)) | (~ spl30_25 | spl30_75)), inference(forward_demodulation, [], [f751, f490])).
fof(f490, plain, ((e0 = op(e2, e1)) | ~ spl30_25), inference(avatar_component_clause, [], [f488])).
fof(f751, plain, (~ (e1 = op(e2, op(e2, e1))) | spl30_75), inference(avatar_component_clause, [], [f750])).
fof(f6351, plain, (~ spl30_18 | ~ spl30_11 | spl30_101), inference(avatar_split_clause, [], [f6350, f928, f428, f458])).
fof(f928, plain, (spl30_101 <=> (e1 = op(op(e3, e1), e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_101])])).
fof(f6350, plain, (~ (e1 = op(e2, e3)) | (~ spl30_11 | spl30_101)), inference(forward_demodulation, [], [f929, f430])).
fof(f430, plain, ((e2 = op(e3, e1)) | ~ spl30_11), inference(avatar_component_clause, [], [f428])).
fof(f929, plain, (~ (e1 = op(op(e3, e1), e3)) | spl30_101), inference(avatar_component_clause, [], [f928])).
fof(f6347, plain, (~ spl30_47 | ~ spl30_54 | spl30_132), inference(avatar_split_clause, [], [f6346, f1085, f611, f581])).
fof(f6346, plain, (~ (e2 = op(e1, e0)) | (~ spl30_54 | spl30_132)), inference(forward_demodulation, [], [f1086, f613])).
fof(f1086, plain, (~ (e2 = op(op(e0, e2), e0)) | spl30_132), inference(avatar_component_clause, [], [f1085])).
fof(f6334, plain, (~ spl30_62 | ~ spl30_54), inference(avatar_split_clause, [], [f6333, f611, f645])).
fof(f645, plain, (spl30_62 <=> (op(e0, e0) = e1)), introduced(avatar_definition, [new_symbols(naming, [spl30_62])])).
fof(f6333, plain, (~ (op(e0, e0) = e1) | ~ spl30_54), inference(forward_demodulation, [], [f188, f613])).
fof(f188, plain, ~ (op(e0, e0) = op(e0, e2)), inference(cnf_transformation, [], [f3])).
fof(f6330, plain, (~ spl30_38 | ~ spl30_54), inference(avatar_split_clause, [], [f6329, f611, f543])).
fof(f543, plain, (spl30_38 <=> (e1 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_38])])).
fof(f6329, plain, (~ (e1 = op(e1, e2)) | ~ spl30_54), inference(forward_demodulation, [], [f175, f613])).
fof(f175, plain, ~ (op(e0, e2) = op(e1, e2)), inference(cnf_transformation, [], [f3])).
fof(f6328, plain, (spl30_37 | ~ spl30_54 | ~ spl30_131), inference(avatar_split_clause, [], [f6288, f1080, f611, f539])).
fof(f1080, plain, (spl30_131 <=> (e0 = op(op(e0, e2), e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_131])])).
fof(f6288, plain, ((e0 = op(e1, e2)) | (~ spl30_54 | ~ spl30_131)), inference(backward_demodulation, [], [f1081, f613])).
fof(f1081, plain, ((e0 = op(op(e0, e2), e2)) | ~ spl30_131), inference(avatar_component_clause, [], [f1080])).
fof(f6309, plain, (~ spl30_19 | ~ spl30_23), inference(avatar_split_clause, [], [f6302, f479, f462])).
fof(f462, plain, (spl30_19 <=> (e2 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_19])])).
fof(f6302, plain, (~ (e2 = op(e2, e3)) | ~ spl30_23), inference(backward_demodulation, [], [f204, f481])).
fof(f481, plain, ((e2 = op(e2, e2)) | ~ spl30_23), inference(avatar_component_clause, [], [f479])).
fof(f204, plain, ~ (op(e2, e2) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f6299, plain, (~ spl30_19 | ~ spl30_51), inference(avatar_split_clause, [], [f6293, f598, f462])).
fof(f6293, plain, (~ (e2 = op(e2, e3)) | ~ spl30_51), inference(backward_demodulation, [], [f182, f600])).
fof(f182, plain, ~ (op(e0, e3) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f6292, plain, (~ spl30_54 | spl30_59 | ~ spl30_90), inference(avatar_contradiction_clause, [], [f6291])).
fof(f6291, plain, ($false | (~ spl30_54 | spl30_59 | ~ spl30_90)), inference(subsumption_resolution, [], [f6287, f633])).
fof(f633, plain, (~ (e2 = op(e0, e1)) | spl30_59), inference(avatar_component_clause, [], [f632])).
fof(f632, plain, (spl30_59 <=> (e2 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_59])])).
fof(f6287, plain, ((e2 = op(e0, e1)) | (~ spl30_54 | ~ spl30_90)), inference(backward_demodulation, [], [f859, f613])).
fof(f859, plain, ((e2 = op(e0, op(e0, e2))) | ~ spl30_90), inference(avatar_component_clause, [], [f857])).
fof(f857, plain, (spl30_90 <=> (e2 = op(e0, op(e0, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl30_90])])).
fof(f6290, plain, (~ spl30_50 | ~ spl30_54), inference(avatar_split_clause, [], [f6284, f611, f594])).
fof(f594, plain, (spl30_50 <=> (e1 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_50])])).
fof(f6284, plain, (~ (e1 = op(e0, e3)) | ~ spl30_54), inference(backward_demodulation, [], [f192, f613])).
fof(f192, plain, ~ (op(e0, e2) = op(e0, e3)), inference(cnf_transformation, [], [f3])).
fof(f6289, plain, (~ spl30_22 | ~ spl30_54), inference(avatar_split_clause, [], [f6283, f611, f475])).
fof(f475, plain, (spl30_22 <=> (e1 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_22])])).
fof(f6283, plain, (~ (e1 = op(e2, e2)) | ~ spl30_54), inference(backward_demodulation, [], [f176, f613])).
fof(f176, plain, ~ (op(e0, e2) = op(e2, e2)), inference(cnf_transformation, [], [f3])).
fof(f6282, plain, (spl30_18 | ~ spl30_11 | ~ spl30_101), inference(avatar_split_clause, [], [f6281, f928, f428, f458])).
fof(f6281, plain, ((e1 = op(e2, e3)) | (~ spl30_11 | ~ spl30_101)), inference(forward_demodulation, [], [f930, f430])).
fof(f930, plain, ((e1 = op(op(e3, e1), e3)) | ~ spl30_101), inference(avatar_component_clause, [], [f928])).
fof(f6280, plain, (spl30_54 | ~ spl30_37 | ~ spl30_120), inference(avatar_split_clause, [], [f6279, f1024, f539, f611])).
fof(f1024, plain, (spl30_120 <=> (e1 = op(op(e1, e2), e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_120])])).
fof(f6279, plain, ((e1 = op(e0, e2)) | (~ spl30_37 | ~ spl30_120)), inference(forward_demodulation, [], [f1025, f541])).
fof(f541, plain, ((e0 = op(e1, e2)) | ~ spl30_37), inference(avatar_component_clause, [], [f539])).
fof(f1025, plain, ((e1 = op(op(e1, e2), e2)) | ~ spl30_120), inference(avatar_component_clause, [], [f1024])).
fof(f6274, plain, (~ spl30_51 | ~ spl30_1 | spl30_145), inference(avatar_split_clause, [], [f6273, f1157, f386, f598])).
fof(f1157, plain, (spl30_145 <=> (e2 = op(op(e3, e3), e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_145])])).
fof(f6273, plain, (~ (e2 = op(e0, e3)) | (~ spl30_1 | spl30_145)), inference(forward_demodulation, [], [f1159, f388])).
fof(f388, plain, ((e0 = op(e3, e3)) | ~ spl30_1), inference(avatar_component_clause, [], [f386])).
fof(f1159, plain, (~ (e2 = op(op(e3, e3), e3)) | spl30_145), inference(avatar_component_clause, [], [f1157])).
fof(f6249, plain, (~ spl30_3 | ~ spl30_11), inference(avatar_split_clause, [], [f6244, f428, f394])).
fof(f394, plain, (spl30_3 <=> (e2 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_3])])).
fof(f6244, plain, (~ (e2 = op(e3, e3)) | ~ spl30_11), inference(backward_demodulation, [], [f209, f430])).
fof(f6240, plain, (~ spl30_3 | ~ spl30_19), inference(avatar_split_clause, [], [f6235, f462, f394])).
fof(f6235, plain, (~ (e2 = op(e3, e3)) | ~ spl30_19), inference(backward_demodulation, [], [f186, f464])).
fof(f464, plain, ((e2 = op(e2, e3)) | ~ spl30_19), inference(avatar_component_clause, [], [f462])).
fof(f6227, plain, (~ spl30_17 | ~ spl30_25), inference(avatar_split_clause, [], [f6222, f488, f454])).
fof(f454, plain, (spl30_17 <=> (e0 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_17])])).
fof(f6222, plain, (~ (e0 = op(e2, e3)) | ~ spl30_25), inference(backward_demodulation, [], [f203, f490])).
fof(f6213, plain, (~ spl30_5 | ~ spl30_37), inference(avatar_split_clause, [], [f6205, f539, f403])).
fof(f6205, plain, (~ (e0 = op(e3, e2)) | ~ spl30_37), inference(backward_demodulation, [], [f179, f541])).
fof(f6196, plain, (~ spl30_34 | ~ spl30_42), inference(avatar_split_clause, [], [f6186, f560, f526])).
fof(f526, plain, (spl30_34 <=> (e1 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_34])])).
fof(f6186, plain, (~ (e1 = op(e1, e3)) | ~ spl30_42), inference(backward_demodulation, [], [f197, f562])).
fof(f6195, plain, (~ spl30_26 | ~ spl30_42), inference(avatar_split_clause, [], [f6183, f560, f492])).
fof(f492, plain, (spl30_26 <=> (e1 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_26])])).
fof(f6183, plain, (~ (e1 = op(e2, e1)) | ~ spl30_42), inference(backward_demodulation, [], [f172, f562])).
fof(f172, plain, ~ (op(e1, e1) = op(e2, e1)), inference(cnf_transformation, [], [f3])).
fof(f6181, plain, (~ spl30_34 | ~ spl30_50), inference(avatar_split_clause, [], [f6173, f594, f526])).
fof(f6173, plain, (~ (e1 = op(e1, e3)) | ~ spl30_50), inference(backward_demodulation, [], [f181, f596])).
fof(f596, plain, ((e1 = op(e0, e3)) | ~ spl30_50), inference(avatar_component_clause, [], [f594])).
fof(f6172, plain, (spl30_31 | ~ spl30_55 | ~ spl30_132), inference(avatar_contradiction_clause, [], [f6171])).
fof(f6171, plain, ($false | (spl30_31 | ~ spl30_55 | ~ spl30_132)), inference(subsumption_resolution, [], [f6169, f514])).
fof(f514, plain, (~ (e2 = op(e2, e0)) | spl30_31), inference(avatar_component_clause, [], [f513])).
fof(f513, plain, (spl30_31 <=> (e2 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_31])])).
fof(f6169, plain, ((e2 = op(e2, e0)) | (~ spl30_55 | ~ spl30_132)), inference(backward_demodulation, [], [f1087, f617])).
fof(f617, plain, ((e2 = op(e0, e2)) | ~ spl30_55), inference(avatar_component_clause, [], [f615])).
fof(f615, plain, (spl30_55 <=> (e2 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_55])])).
fof(f6170, plain, (~ spl30_23 | ~ spl30_55), inference(avatar_split_clause, [], [f6163, f615, f479])).
fof(f6163, plain, (~ (e2 = op(e2, e2)) | ~ spl30_55), inference(backward_demodulation, [], [f176, f617])).
fof(f6161, plain, (~ spl30_50 | ~ spl30_60 | spl30_91), inference(avatar_split_clause, [], [f6156, f862, f636, f594])).
fof(f862, plain, (spl30_91 <=> (e1 = op(e0, op(e0, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl30_91])])).
fof(f6156, plain, (~ (e1 = op(e0, e3)) | (~ spl30_60 | spl30_91)), inference(backward_demodulation, [], [f863, f638])).
fof(f863, plain, (~ (e1 = op(e0, op(e0, e1))) | spl30_91), inference(avatar_component_clause, [], [f862])).
fof(f6160, plain, (~ spl30_52 | ~ spl30_60), inference(avatar_split_clause, [], [f6155, f636, f602])).
fof(f602, plain, (spl30_52 <=> (e3 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_52])])).
fof(f6155, plain, (~ (e3 = op(e0, e3)) | ~ spl30_60), inference(backward_demodulation, [], [f191, f638])).
fof(f6159, plain, (~ spl30_12 | ~ spl30_60), inference(avatar_split_clause, [], [f6153, f636, f432])).
fof(f432, plain, (spl30_12 <=> (e3 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_12])])).
fof(f6153, plain, (~ (e3 = op(e3, e1)) | ~ spl30_60), inference(backward_demodulation, [], [f171, f638])).
fof(f171, plain, ~ (op(e0, e1) = op(e3, e1)), inference(cnf_transformation, [], [f3])).
fof(f6150, plain, (~ spl30_17 | ~ spl30_32 | spl30_76), inference(avatar_split_clause, [], [f6149, f755, f517, f454])).
fof(f755, plain, (spl30_76 <=> (e0 = op(e2, op(e2, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl30_76])])).
fof(f6149, plain, (~ (e0 = op(e2, e3)) | (~ spl30_32 | spl30_76)), inference(forward_demodulation, [], [f756, f519])).
fof(f756, plain, (~ (e0 = op(e2, op(e2, e0))) | spl30_76), inference(avatar_component_clause, [], [f755])).
fof(f6148, plain, (spl30_37 | ~ spl30_47 | ~ spl30_84), inference(avatar_split_clause, [], [f6147, f811, f581, f539])).
fof(f811, plain, (spl30_84 <=> (e0 = op(e1, op(e1, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl30_84])])).
fof(f6147, plain, ((e0 = op(e1, e2)) | (~ spl30_47 | ~ spl30_84)), inference(forward_demodulation, [], [f813, f583])).
fof(f813, plain, ((e0 = op(e1, op(e1, e0))) | ~ spl30_84), inference(avatar_component_clause, [], [f811])).
fof(f6144, plain, (spl30_25 | ~ spl30_47 | ~ spl30_126), inference(avatar_split_clause, [], [f6143, f1055, f581, f488])).
fof(f1055, plain, (spl30_126 <=> (e0 = op(op(e1, e0), e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_126])])).
fof(f6143, plain, ((e0 = op(e2, e1)) | (~ spl30_47 | ~ spl30_126)), inference(forward_demodulation, [], [f1057, f583])).
fof(f1057, plain, ((e0 = op(op(e1, e0), e1)) | ~ spl30_126), inference(avatar_component_clause, [], [f1055])).
fof(f6115, plain, (~ spl30_1 | ~ spl30_17), inference(avatar_split_clause, [], [f6108, f454, f386])).
fof(f6108, plain, (~ (e0 = op(e3, e3)) | ~ spl30_17), inference(backward_demodulation, [], [f186, f456])).
fof(f456, plain, ((e0 = op(e2, e3)) | ~ spl30_17), inference(avatar_component_clause, [], [f454])).
fof(f6105, plain, (~ spl30_7 | ~ spl30_23), inference(avatar_split_clause, [], [f6098, f479, f411])).
fof(f411, plain, (spl30_7 <=> (e2 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_7])])).
fof(f6098, plain, (~ (e2 = op(e3, e2)) | ~ spl30_23), inference(backward_demodulation, [], [f180, f481])).
fof(f6097, plain, (~ spl30_18 | ~ spl30_26), inference(avatar_split_clause, [], [f6093, f492, f458])).
fof(f6093, plain, (~ (e1 = op(e2, e3)) | ~ spl30_26), inference(backward_demodulation, [], [f203, f494])).
fof(f494, plain, ((e1 = op(e2, e1)) | ~ spl30_26), inference(avatar_component_clause, [], [f492])).
fof(f6091, plain, (spl30_5 | ~ spl30_32 | ~ spl30_115), inference(avatar_split_clause, [], [f6087, f999, f517, f403])).
fof(f6087, plain, ((e0 = op(e3, e2)) | (~ spl30_32 | ~ spl30_115)), inference(backward_demodulation, [], [f1001, f519])).
fof(f1001, plain, ((e0 = op(op(e2, e0), e2)) | ~ spl30_115), inference(avatar_component_clause, [], [f999])).
fof(f6089, plain, (spl30_17 | ~ spl30_32 | ~ spl30_76), inference(avatar_split_clause, [], [f6085, f755, f517, f454])).
fof(f6085, plain, ((e0 = op(e2, e3)) | (~ spl30_32 | ~ spl30_76)), inference(backward_demodulation, [], [f757, f519])).
fof(f757, plain, ((e0 = op(e2, op(e2, e0))) | ~ spl30_76), inference(avatar_component_clause, [], [f755])).
fof(f6088, plain, (~ spl30_24 | ~ spl30_32), inference(avatar_split_clause, [], [f6083, f517, f483])).
fof(f6083, plain, (~ (e3 = op(e2, e2)) | ~ spl30_32), inference(backward_demodulation, [], [f200, f519])).
fof(f200, plain, ~ (op(e2, e0) = op(e2, e2)), inference(cnf_transformation, [], [f3])).
fof(f6079, plain, (~ spl30_18 | ~ spl30_34), inference(avatar_split_clause, [], [f6072, f526, f458])).
fof(f6072, plain, (~ (e1 = op(e2, e3)) | ~ spl30_34), inference(backward_demodulation, [], [f184, f528])).
fof(f528, plain, ((e1 = op(e1, e3)) | ~ spl30_34), inference(avatar_component_clause, [], [f526])).
fof(f184, plain, ~ (op(e1, e3) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f6071, plain, (spl30_6 | ~ spl30_40 | ~ spl30_120), inference(avatar_contradiction_clause, [], [f6070])).
fof(f6070, plain, ($false | (spl30_6 | ~ spl30_40 | ~ spl30_120)), inference(subsumption_resolution, [], [f6066, f408])).
fof(f408, plain, (~ (e1 = op(e3, e2)) | spl30_6), inference(avatar_component_clause, [], [f407])).
fof(f407, plain, (spl30_6 <=> (e1 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_6])])).
fof(f6066, plain, ((e1 = op(e3, e2)) | (~ spl30_40 | ~ spl30_120)), inference(backward_demodulation, [], [f1025, f553])).
fof(f6061, plain, (~ spl30_37 | ~ spl30_41), inference(avatar_split_clause, [], [f6052, f556, f539])).
fof(f556, plain, (spl30_41 <=> (e0 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_41])])).
fof(f6052, plain, (~ (e0 = op(e1, e2)) | ~ spl30_41), inference(backward_demodulation, [], [f196, f558])).
fof(f558, plain, ((e0 = op(e1, e1)) | ~ spl30_41), inference(avatar_component_clause, [], [f556])).
fof(f196, plain, ~ (op(e1, e1) = op(e1, e2)), inference(cnf_transformation, [], [f3])).
fof(f6060, plain, (~ spl30_25 | ~ spl30_41), inference(avatar_split_clause, [], [f6051, f556, f488])).
fof(f6051, plain, (~ (e0 = op(e2, e1)) | ~ spl30_41), inference(backward_demodulation, [], [f172, f558])).
fof(f6049, plain, (~ spl30_37 | ~ spl30_47 | spl30_84), inference(avatar_split_clause, [], [f6044, f811, f581, f539])).
fof(f6044, plain, (~ (e0 = op(e1, e2)) | (~ spl30_47 | spl30_84)), inference(backward_demodulation, [], [f812, f583])).
fof(f812, plain, (~ (e0 = op(e1, op(e1, e0))) | spl30_84), inference(avatar_component_clause, [], [f811])).
fof(f6048, plain, (~ spl30_35 | ~ spl30_47), inference(avatar_split_clause, [], [f6042, f581, f530])).
fof(f6042, plain, (~ (e2 = op(e1, e3)) | ~ spl30_47), inference(backward_demodulation, [], [f195, f583])).
fof(f195, plain, ~ (op(e1, e0) = op(e1, e3)), inference(cnf_transformation, [], [f3])).
fof(f6047, plain, (~ spl30_31 | ~ spl30_47), inference(avatar_split_clause, [], [f6039, f581, f513])).
fof(f6039, plain, (~ (e2 = op(e2, e0)) | ~ spl30_47), inference(backward_demodulation, [], [f166, f583])).
fof(f6038, plain, (~ spl30_54 | ~ spl30_59 | spl30_90), inference(avatar_contradiction_clause, [], [f6037])).
fof(f6037, plain, ($false | (~ spl30_54 | ~ spl30_59 | spl30_90)), inference(subsumption_resolution, [], [f6036, f634])).
fof(f634, plain, ((e2 = op(e0, e1)) | ~ spl30_59), inference(avatar_component_clause, [], [f632])).
fof(f6036, plain, (~ (e2 = op(e0, e1)) | (~ spl30_54 | spl30_90)), inference(forward_demodulation, [], [f858, f613])).
fof(f858, plain, (~ (e2 = op(e0, op(e0, e2))) | spl30_90), inference(avatar_component_clause, [], [f857])).
fof(f6035, plain, (~ spl30_48 | ~ spl30_14 | spl30_103), inference(avatar_split_clause, [], [f6034, f938, f441, f585])).
fof(f938, plain, (spl30_103 <=> (e3 = op(op(e3, e0), e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_103])])).
fof(f6034, plain, (~ (e3 = op(e1, e0)) | (~ spl30_14 | spl30_103)), inference(forward_demodulation, [], [f940, f443])).
fof(f940, plain, (~ (e3 = op(op(e3, e0), e0)) | spl30_103), inference(avatar_component_clause, [], [f938])).
fof(f6031, plain, (~ spl30_37 | ~ spl30_54 | spl30_131), inference(avatar_split_clause, [], [f6030, f1080, f611, f539])).
fof(f6030, plain, (~ (e0 = op(e1, e2)) | (~ spl30_54 | spl30_131)), inference(forward_demodulation, [], [f1082, f613])).
fof(f1082, plain, (~ (e0 = op(op(e0, e2), e2)) | spl30_131), inference(avatar_component_clause, [], [f1080])).
fof(f6027, plain, (~ spl30_25 | ~ spl30_59 | spl30_134), inference(avatar_split_clause, [], [f6026, f1095, f632, f488])).
fof(f6026, plain, (~ (e0 = op(e2, e1)) | (~ spl30_59 | spl30_134)), inference(forward_demodulation, [], [f1097, f634])).
fof(f1097, plain, (~ (e0 = op(op(e0, e1), e1)) | spl30_134), inference(avatar_component_clause, [], [f1095])).
fof(f6025, plain, (~ spl30_25 | spl30_143 | ~ spl30_144), inference(avatar_split_clause, [], [f6023, f1152, f1148, f488])).
fof(f1148, plain, (spl30_143 <=> (e0 = op(op(op(e1, e1), e1), e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_143])])).
fof(f1152, plain, (spl30_144 <=> (e2 = op(op(e1, e1), e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_144])])).
fof(f6023, plain, (~ (e0 = op(e2, e1)) | (spl30_143 | ~ spl30_144)), inference(backward_demodulation, [], [f1150, f1153])).
fof(f1153, plain, ((e2 = op(op(e1, e1), e1)) | ~ spl30_144), inference(avatar_component_clause, [], [f1152])).
fof(f1150, plain, (~ (e0 = op(op(op(e1, e1), e1), e1)) | spl30_143), inference(avatar_component_clause, [], [f1148])).
fof(f6014, plain, (~ spl30_44 | ~ spl30_12), inference(avatar_split_clause, [], [f6013, f432, f568])).
fof(f568, plain, (spl30_44 <=> (e3 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_44])])).
fof(f6013, plain, (~ (e3 = op(e1, e1)) | ~ spl30_12), inference(forward_demodulation, [], [f173, f434])).
fof(f434, plain, ((e3 = op(e3, e1)) | ~ spl30_12), inference(avatar_component_clause, [], [f432])).
fof(f5989, plain, (~ spl30_8 | ~ spl30_12), inference(avatar_split_clause, [], [f5984, f432, f415])).
fof(f5984, plain, (~ (e3 = op(e3, e2)) | ~ spl30_12), inference(backward_demodulation, [], [f208, f434])).
fof(f5983, plain, (~ spl30_9 | ~ spl30_14 | spl30_69), inference(avatar_split_clause, [], [f5979, f709, f441, f420])).
fof(f5979, plain, (~ (e0 = op(e3, e1)) | (~ spl30_14 | spl30_69)), inference(backward_demodulation, [], [f710, f443])).
fof(f710, plain, (~ (e0 = op(e3, op(e3, e0))) | spl30_69), inference(avatar_component_clause, [], [f709])).
fof(f5982, plain, (~ spl30_2 | ~ spl30_14), inference(avatar_split_clause, [], [f5976, f441, f390])).
fof(f5976, plain, (~ (e1 = op(e3, e3)) | ~ spl30_14), inference(backward_demodulation, [], [f207, f443])).
fof(f207, plain, ~ (op(e3, e0) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f5964, plain, (spl30_19 | ~ spl30_24 | ~ spl30_74), inference(avatar_contradiction_clause, [], [f5963])).
fof(f5963, plain, ($false | (spl30_19 | ~ spl30_24 | ~ spl30_74)), inference(subsumption_resolution, [], [f5957, f463])).
fof(f463, plain, (~ (e2 = op(e2, e3)) | spl30_19), inference(avatar_component_clause, [], [f462])).
fof(f5957, plain, ((e2 = op(e2, e3)) | (~ spl30_24 | ~ spl30_74)), inference(backward_demodulation, [], [f747, f485])).
fof(f485, plain, ((e3 = op(e2, e2)) | ~ spl30_24), inference(avatar_component_clause, [], [f483])).
fof(f747, plain, ((e2 = op(e2, op(e2, e2))) | ~ spl30_74), inference(avatar_component_clause, [], [f745])).
fof(f745, plain, (spl30_74 <=> (e2 = op(e2, op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl30_74])])).
fof(f5962, plain, (~ spl30_8 | ~ spl30_24), inference(avatar_split_clause, [], [f5955, f483, f415])).
fof(f5955, plain, (~ (e3 = op(e3, e2)) | ~ spl30_24), inference(backward_demodulation, [], [f180, f485])).
fof(f5954, plain, (~ spl30_25 | ~ spl30_59 | spl30_111), inference(avatar_contradiction_clause, [], [f5953])).
fof(f5953, plain, ($false | (~ spl30_25 | ~ spl30_59 | spl30_111)), inference(subsumption_resolution, [], [f5949, f634])).
fof(f5949, plain, (~ (e2 = op(e0, e1)) | (~ spl30_25 | spl30_111)), inference(backward_demodulation, [], [f981, f490])).
fof(f5951, plain, (~ spl30_9 | ~ spl30_25), inference(avatar_split_clause, [], [f5942, f488, f420])).
fof(f5942, plain, (~ (e0 = op(e3, e1)) | ~ spl30_25), inference(backward_demodulation, [], [f174, f490])).
fof(f174, plain, ~ (op(e2, e1) = op(e3, e1)), inference(cnf_transformation, [], [f3])).
fof(f5941, plain, (~ spl30_23 | ~ spl30_31), inference(avatar_split_clause, [], [f5934, f513, f479])).
fof(f5934, plain, (~ (e2 = op(e2, e2)) | ~ spl30_31), inference(backward_demodulation, [], [f200, f515])).
fof(f515, plain, ((e2 = op(e2, e0)) | ~ spl30_31), inference(avatar_component_clause, [], [f513])).
fof(f5940, plain, (~ spl30_15 | ~ spl30_31), inference(avatar_split_clause, [], [f5932, f513, f445])).
fof(f5932, plain, (~ (e2 = op(e3, e0)) | ~ spl30_31), inference(backward_demodulation, [], [f168, f515])).
fof(f5931, plain, (spl30_18 | ~ spl30_35 | ~ spl30_117), inference(avatar_split_clause, [], [f5930, f1009, f530, f458])).
fof(f5930, plain, ((e1 = op(e2, e3)) | (~ spl30_35 | ~ spl30_117)), inference(forward_demodulation, [], [f1010, f532])).
fof(f1010, plain, ((e1 = op(op(e1, e3), e3)) | ~ spl30_117), inference(avatar_component_clause, [], [f1009])).
fof(f5927, plain, (spl30_14 | ~ spl30_48 | ~ spl30_125), inference(avatar_split_clause, [], [f5926, f1050, f585, f441])).
fof(f1050, plain, (spl30_125 <=> (e1 = op(op(e1, e0), e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_125])])).
fof(f5926, plain, ((e1 = op(e3, e0)) | (~ spl30_48 | ~ spl30_125)), inference(forward_demodulation, [], [f1051, f587])).
fof(f587, plain, ((e3 = op(e1, e0)) | ~ spl30_48), inference(avatar_component_clause, [], [f585])).
fof(f1051, plain, ((e1 = op(op(e1, e0), e0)) | ~ spl30_125), inference(avatar_component_clause, [], [f1050])).
fof(f5921, plain, (spl30_25 | ~ spl30_59 | ~ spl30_134), inference(avatar_split_clause, [], [f5920, f1095, f632, f488])).
fof(f5920, plain, ((e0 = op(e2, e1)) | (~ spl30_59 | ~ spl30_134)), inference(forward_demodulation, [], [f1096, f634])).
fof(f5907, plain, (~ spl30_32 | ~ spl30_48), inference(avatar_split_clause, [], [f5906, f585, f517])).
fof(f5906, plain, (~ (e3 = op(e2, e0)) | ~ spl30_48), inference(forward_demodulation, [], [f166, f587])).
fof(f5900, plain, (~ spl30_19 | ~ spl30_35), inference(avatar_split_clause, [], [f5899, f530, f462])).
fof(f5899, plain, (~ (e2 = op(e2, e3)) | ~ spl30_35), inference(forward_demodulation, [], [f184, f532])).
fof(f5898, plain, (~ spl30_20 | ~ spl30_52), inference(avatar_split_clause, [], [f5897, f602, f466])).
fof(f5897, plain, (~ (e3 = op(e2, e3)) | ~ spl30_52), inference(forward_demodulation, [], [f182, f604])).
fof(f604, plain, ((e3 = op(e0, e3)) | ~ spl30_52), inference(avatar_component_clause, [], [f602])).
fof(f5893, plain, (~ spl30_11 | ~ spl30_59), inference(avatar_split_clause, [], [f5892, f632, f428])).
fof(f5892, plain, (~ (e2 = op(e3, e1)) | ~ spl30_59), inference(forward_demodulation, [], [f171, f634])).
fof(f5891, plain, (~ spl30_2 | ~ spl30_35 | spl30_145), inference(avatar_contradiction_clause, [], [f5890])).
fof(f5890, plain, ($false | (~ spl30_2 | ~ spl30_35 | spl30_145)), inference(subsumption_resolution, [], [f5887, f532])).
fof(f5887, plain, (~ (e2 = op(e1, e3)) | (~ spl30_2 | spl30_145)), inference(backward_demodulation, [], [f1159, f392])).
fof(f392, plain, ((e1 = op(e3, e3)) | ~ spl30_2), inference(avatar_component_clause, [], [f390])).
fof(f5883, plain, (~ spl30_4 | ~ spl30_8), inference(avatar_split_clause, [], [f5879, f415, f398])).
fof(f398, plain, (spl30_4 <=> (e3 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_4])])).
fof(f5879, plain, (~ (e3 = op(e3, e3)) | ~ spl30_8), inference(backward_demodulation, [], [f210, f417])).
fof(f417, plain, ((e3 = op(e3, e2)) | ~ spl30_8), inference(avatar_component_clause, [], [f415])).
fof(f5871, plain, (~ spl30_42 | ~ spl30_146), inference(avatar_contradiction_clause, [], [f5870])).
fof(f5870, plain, ($false | (~ spl30_42 | ~ spl30_146)), inference(subsumption_resolution, [], [f5869, f215])).
fof(f215, plain, ~ (e1 = e3), inference(cnf_transformation, [], [f4])).
fof(f4, plain, (~ (e2 = e3) & ~ (e1 = e3) & ~ (e1 = e2) & ~ (e0 = e3) & ~ (e0 = e2) & ~ (e0 = e1)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG164+1.p', ax4)).
fof(f5869, plain, ((e1 = e3) | (~ spl30_42 | ~ spl30_146)), inference(forward_demodulation, [], [f5866, f562])).
fof(f5866, plain, ((e3 = op(e1, e1)) | (~ spl30_42 | ~ spl30_146)), inference(backward_demodulation, [], [f1163, f562])).
fof(f1163, plain, ((e3 = op(op(e1, e1), e1)) | ~ spl30_146), inference(avatar_component_clause, [], [f1162])).
fof(f1162, plain, (spl30_146 <=> (e3 = op(op(e1, e1), e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_146])])).
fof(f5860, plain, (~ spl30_4 | ~ spl30_52), inference(avatar_split_clause, [], [f5856, f602, f398])).
fof(f5856, plain, (~ (e3 = op(e3, e3)) | ~ spl30_52), inference(backward_demodulation, [], [f183, f604])).
fof(f5854, plain, (~ spl30_6 | ~ spl30_54), inference(avatar_split_clause, [], [f5847, f611, f407])).
fof(f5847, plain, (~ (e1 = op(e3, e2)) | ~ spl30_54), inference(backward_demodulation, [], [f177, f613])).
fof(f177, plain, ~ (op(e0, e2) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f5844, plain, (spl30_54 | ~ spl30_59 | ~ spl30_91), inference(avatar_split_clause, [], [f5839, f862, f632, f611])).
fof(f5839, plain, ((e1 = op(e0, e2)) | (~ spl30_59 | ~ spl30_91)), inference(backward_demodulation, [], [f864, f634])).
fof(f864, plain, ((e1 = op(e0, op(e0, e1))) | ~ spl30_91), inference(avatar_component_clause, [], [f862])).
fof(f5843, plain, (~ spl30_51 | ~ spl30_59), inference(avatar_split_clause, [], [f5837, f632, f598])).
fof(f5837, plain, (~ (e2 = op(e0, e3)) | ~ spl30_59), inference(backward_demodulation, [], [f191, f634])).
fof(f5842, plain, (~ spl30_43 | ~ spl30_59), inference(avatar_split_clause, [], [f5835, f632, f564])).
fof(f5835, plain, (~ (e2 = op(e1, e1)) | ~ spl30_59), inference(backward_demodulation, [], [f169, f634])).
fof(f169, plain, ~ (op(e0, e1) = op(e1, e1)), inference(cnf_transformation, [], [f3])).
fof(f5834, plain, (~ spl30_51 | ~ spl30_17 | spl30_106), inference(avatar_split_clause, [], [f5833, f953, f454, f598])).
fof(f953, plain, (spl30_106 <=> (e2 = op(op(e2, e3), e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_106])])).
fof(f5833, plain, (~ (e2 = op(e0, e3)) | (~ spl30_17 | spl30_106)), inference(forward_demodulation, [], [f955, f456])).
fof(f955, plain, (~ (e2 = op(op(e2, e3), e3)) | spl30_106), inference(avatar_component_clause, [], [f953])).
fof(f5830, plain, (~ spl30_6 | ~ spl30_28 | spl30_112), inference(avatar_split_clause, [], [f5829, f984, f500, f407])).
fof(f984, plain, (spl30_112 <=> (e1 = op(op(e2, e1), e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_112])])).
fof(f5829, plain, (~ (e1 = op(e3, e2)) | (~ spl30_28 | spl30_112)), inference(forward_demodulation, [], [f985, f502])).
fof(f985, plain, (~ (e1 = op(op(e2, e1), e2)) | spl30_112), inference(avatar_component_clause, [], [f984])).
fof(f5828, plain, (spl30_54 | ~ spl30_37 | ~ spl30_120), inference(avatar_split_clause, [], [f5827, f1024, f539, f611])).
fof(f5827, plain, ((e1 = op(e0, e2)) | (~ spl30_37 | ~ spl30_120)), inference(forward_demodulation, [], [f1025, f541])).
fof(f5826, plain, (spl30_59 | ~ spl30_37 | ~ spl30_121), inference(avatar_split_clause, [], [f5825, f1029, f539, f632])).
fof(f5825, plain, ((e2 = op(e0, e1)) | (~ spl30_37 | ~ spl30_121)), inference(forward_demodulation, [], [f1031, f541])).
fof(f1031, plain, ((e2 = op(op(e1, e2), e1)) | ~ spl30_121), inference(avatar_component_clause, [], [f1029])).
fof(f5808, plain, (~ spl30_7 | ~ spl30_15), inference(avatar_split_clause, [], [f5807, f445, f411])).
fof(f5807, plain, (~ (e2 = op(e3, e2)) | ~ spl30_15), inference(forward_demodulation, [], [f206, f447])).
fof(f5798, plain, (~ spl30_9 | ~ spl30_28 | ~ spl30_43 | spl30_143), inference(avatar_contradiction_clause, [], [f5797])).
fof(f5797, plain, ($false | (~ spl30_9 | ~ spl30_28 | ~ spl30_43 | spl30_143)), inference(subsumption_resolution, [], [f5796, f422])).
fof(f5796, plain, (~ (e0 = op(e3, e1)) | (~ spl30_28 | ~ spl30_43 | spl30_143)), inference(forward_demodulation, [], [f5795, f502])).
fof(f5795, plain, (~ (e0 = op(op(e2, e1), e1)) | (~ spl30_43 | spl30_143)), inference(forward_demodulation, [], [f1150, f566])).
fof(f5764, plain, (~ spl30_28 | ~ spl30_43 | spl30_146), inference(avatar_contradiction_clause, [], [f5763])).
fof(f5763, plain, ($false | (~ spl30_28 | ~ spl30_43 | spl30_146)), inference(subsumption_resolution, [], [f5756, f502])).
fof(f5756, plain, (~ (e3 = op(e2, e1)) | (~ spl30_43 | spl30_146)), inference(backward_demodulation, [], [f1164, f566])).
fof(f1164, plain, (~ (e3 = op(op(e1, e1), e1)) | spl30_146), inference(avatar_component_clause, [], [f1162])).
fof(f5751, plain, (~ spl30_30 | ~ spl30_51 | ~ spl30_129), inference(avatar_contradiction_clause, [], [f5750])).
fof(f5750, plain, ($false | (~ spl30_30 | ~ spl30_51 | ~ spl30_129)), inference(subsumption_resolution, [], [f5749, f215])).
fof(f5749, plain, ((e1 = e3) | (~ spl30_30 | ~ spl30_51 | ~ spl30_129)), inference(forward_demodulation, [], [f5747, f511])).
fof(f5747, plain, ((e3 = op(e2, e0)) | (~ spl30_51 | ~ spl30_129)), inference(backward_demodulation, [], [f1072, f600])).
fof(f5742, plain, (~ spl30_50 | ~ spl30_58), inference(avatar_split_clause, [], [f5735, f628, f594])).
fof(f628, plain, (spl30_58 <=> (e1 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_58])])).
fof(f5735, plain, (~ (e1 = op(e0, e3)) | ~ spl30_58), inference(backward_demodulation, [], [f191, f630])).
fof(f630, plain, ((e1 = op(e0, e1)) | ~ spl30_58), inference(avatar_component_clause, [], [f628])).
fof(f5741, plain, (~ spl30_42 | ~ spl30_58), inference(avatar_split_clause, [], [f5733, f628, f560])).
fof(f5733, plain, (~ (e1 = op(e1, e1)) | ~ spl30_58), inference(backward_demodulation, [], [f169, f630])).
fof(f5721, plain, (~ spl30_13 | ~ spl30_61), inference(avatar_split_clause, [], [f5709, f641, f437])).
fof(f437, plain, (spl30_13 <=> (e0 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_13])])).
fof(f5709, plain, (~ (e0 = op(e3, e0)) | ~ spl30_61), inference(backward_demodulation, [], [f165, f643])).
fof(f5708, plain, (~ spl30_11 | ~ spl30_6 | spl30_67), inference(avatar_split_clause, [], [f5707, f699, f407, f428])).
fof(f5707, plain, (~ (e2 = op(e3, e1)) | (~ spl30_6 | spl30_67)), inference(forward_demodulation, [], [f700, f409])).
fof(f409, plain, ((e1 = op(e3, e2)) | ~ spl30_6), inference(avatar_component_clause, [], [f407])).
fof(f5706, plain, (spl30_51 | ~ spl30_56 | ~ spl30_90), inference(avatar_split_clause, [], [f5705, f857, f619, f598])).
fof(f619, plain, (spl30_56 <=> (e3 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_56])])).
fof(f5705, plain, ((e2 = op(e0, e3)) | (~ spl30_56 | ~ spl30_90)), inference(forward_demodulation, [], [f859, f621])).
fof(f621, plain, ((e3 = op(e0, e2)) | ~ spl30_56), inference(avatar_component_clause, [], [f619])).
fof(f5698, plain, (spl30_9 | ~ spl30_48 | ~ spl30_126), inference(avatar_split_clause, [], [f5697, f1055, f585, f420])).
fof(f5697, plain, ((e0 = op(e3, e1)) | (~ spl30_48 | ~ spl30_126)), inference(forward_demodulation, [], [f1057, f587])).
fof(f5689, plain, (~ spl30_59 | ~ spl30_37 | spl30_121), inference(avatar_split_clause, [], [f5653, f1029, f539, f632])).
fof(f5653, plain, (~ (e2 = op(e0, e1)) | (~ spl30_37 | spl30_121)), inference(backward_demodulation, [], [f1030, f541])).
fof(f5684, plain, (~ spl30_16 | ~ spl30_48), inference(avatar_split_clause, [], [f5683, f585, f449])).
fof(f5683, plain, (~ (e3 = op(e3, e0)) | ~ spl30_48), inference(forward_demodulation, [], [f167, f587])).
fof(f5682, plain, (~ spl30_15 | ~ spl30_56 | spl30_132), inference(avatar_split_clause, [], [f5607, f1085, f619, f445])).
fof(f5607, plain, (~ (e2 = op(e3, e0)) | (~ spl30_56 | spl30_132)), inference(backward_demodulation, [], [f1086, f621])).
fof(f5681, plain, (~ spl30_14 | ~ spl30_48 | spl30_125), inference(avatar_split_clause, [], [f5624, f1050, f585, f441])).
fof(f5624, plain, (~ (e1 = op(e3, e0)) | (~ spl30_48 | spl30_125)), inference(backward_demodulation, [], [f1052, f587])).
fof(f1052, plain, (~ (e1 = op(op(e1, e0), e0)) | spl30_125), inference(avatar_component_clause, [], [f1050])).
fof(f5654, plain, (~ spl30_21 | ~ spl30_37), inference(avatar_split_clause, [], [f5647, f539, f471])).
fof(f5647, plain, (~ (e0 = op(e2, e2)) | ~ spl30_37), inference(backward_demodulation, [], [f178, f541])).
fof(f5626, plain, (~ spl30_40 | ~ spl30_48), inference(avatar_split_clause, [], [f5621, f585, f551])).
fof(f5621, plain, (~ (e3 = op(e1, e2)) | ~ spl30_48), inference(backward_demodulation, [], [f194, f587])).
fof(f5619, plain, (spl30_48 | ~ spl30_50 | ~ spl30_129), inference(avatar_split_clause, [], [f5616, f1070, f594, f585])).
fof(f5616, plain, ((e3 = op(e1, e0)) | (~ spl30_50 | ~ spl30_129)), inference(backward_demodulation, [], [f1072, f596])).
fof(f5608, plain, (~ spl30_40 | ~ spl30_56), inference(avatar_split_clause, [], [f5602, f619, f551])).
fof(f5602, plain, (~ (e3 = op(e1, e2)) | ~ spl30_56), inference(backward_demodulation, [], [f175, f621])).
fof(f5600, plain, (~ spl30_41 | ~ spl30_57), inference(avatar_split_clause, [], [f5591, f624, f556])).
fof(f5591, plain, (~ (e0 = op(e1, e1)) | ~ spl30_57), inference(backward_demodulation, [], [f169, f626])).
fof(f5590, plain, (spl30_48 | ~ spl30_30 | ~ spl30_63 | ~ spl30_162), inference(avatar_split_clause, [], [f5589, f1242, f649, f509, f585])).
fof(f1242, plain, (spl30_162 <=> (e3 = op(op(op(e0, e0), e0), e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_162])])).
fof(f5589, plain, ((e3 = op(e1, e0)) | (~ spl30_30 | ~ spl30_63 | ~ spl30_162)), inference(forward_demodulation, [], [f5580, f511])).
fof(f5580, plain, ((e3 = op(op(e2, e0), e0)) | (~ spl30_63 | ~ spl30_162)), inference(backward_demodulation, [], [f1243, f651])).
fof(f1243, plain, ((e3 = op(op(op(e0, e0), e0), e0)) | ~ spl30_162), inference(avatar_component_clause, [], [f1242])).
fof(f5588, plain, (~ spl30_30 | ~ spl30_63 | spl30_159), inference(avatar_contradiction_clause, [], [f5587])).
fof(f5587, plain, ($false | (~ spl30_30 | ~ spl30_63 | spl30_159)), inference(subsumption_resolution, [], [f5579, f511])).
fof(f5579, plain, (~ (e1 = op(e2, e0)) | (~ spl30_63 | spl30_159)), inference(backward_demodulation, [], [f1226, f651])).
fof(f1226, plain, (~ (e1 = op(op(e0, e0), e0)) | spl30_159), inference(avatar_component_clause, [], [f1224])).
fof(f5582, plain, (~ spl30_55 | ~ spl30_63), inference(avatar_split_clause, [], [f5572, f649, f615])).
fof(f5572, plain, (~ (e2 = op(e0, e2)) | ~ spl30_63), inference(backward_demodulation, [], [f188, f651])).
fof(f5569, plain, (~ spl30_40 | ~ spl30_6 | spl30_97), inference(avatar_split_clause, [], [f5568, f908, f407, f551])).
fof(f908, plain, (spl30_97 <=> (e3 = op(op(e3, e2), e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_97])])).
fof(f5568, plain, (~ (e3 = op(e1, e2)) | (~ spl30_6 | spl30_97)), inference(forward_demodulation, [], [f910, f409])).
fof(f910, plain, (~ (e3 = op(op(e3, e2), e2)) | spl30_97), inference(avatar_component_clause, [], [f908])).
fof(f5567, plain, (spl30_35 | ~ spl30_6 | ~ spl30_98), inference(avatar_split_clause, [], [f5566, f913, f407, f530])).
fof(f5566, plain, ((e2 = op(e1, e3)) | (~ spl30_6 | ~ spl30_98)), inference(forward_demodulation, [], [f915, f409])).
fof(f5565, plain, (~ spl30_64 | ~ spl30_13 | spl30_103), inference(avatar_split_clause, [], [f5564, f938, f437, f653])).
fof(f5564, plain, (~ (op(e0, e0) = e3) | (~ spl30_13 | spl30_103)), inference(forward_demodulation, [], [f940, f439])).
fof(f439, plain, ((e0 = op(e3, e0)) | ~ spl30_13), inference(avatar_component_clause, [], [f437])).
fof(f5563, plain, (~ spl30_47 | ~ spl30_30 | spl30_114), inference(avatar_split_clause, [], [f5562, f994, f509, f581])).
fof(f5562, plain, (~ (e2 = op(e1, e0)) | (~ spl30_30 | spl30_114)), inference(forward_demodulation, [], [f996, f511])).
fof(f5561, plain, (spl30_37 | ~ spl30_30 | ~ spl30_115), inference(avatar_split_clause, [], [f5560, f999, f509, f539])).
fof(f5560, plain, ((e0 = op(e1, e2)) | (~ spl30_30 | ~ spl30_115)), inference(forward_demodulation, [], [f1001, f511])).
fof(f5540, plain, (~ spl30_1 | ~ spl30_13), inference(avatar_split_clause, [], [f5531, f437, f386])).
fof(f5531, plain, (~ (e0 = op(e3, e3)) | ~ spl30_13), inference(backward_demodulation, [], [f207, f439])).
fof(f5530, plain, (~ spl30_49 | spl30_128), inference(avatar_contradiction_clause, [], [f5529])).
fof(f5529, plain, ($false | (~ spl30_49 | spl30_128)), inference(subsumption_resolution, [], [f5524, f592])).
fof(f5524, plain, (~ (e0 = op(e0, e3)) | (~ spl30_49 | spl30_128)), inference(backward_demodulation, [], [f1067, f592])).
fof(f1067, plain, (~ (e0 = op(op(e0, e3), e3)) | spl30_128), inference(avatar_component_clause, [], [f1065])).
fof(f1065, plain, (spl30_128 <=> (e0 = op(op(e0, e3), e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_128])])).
fof(f5528, plain, (~ spl30_49 | ~ spl30_64 | spl30_89), inference(avatar_contradiction_clause, [], [f5527])).
fof(f5527, plain, ($false | (~ spl30_49 | ~ spl30_64 | spl30_89)), inference(subsumption_resolution, [], [f5523, f655])).
fof(f5523, plain, (~ (op(e0, e0) = e3) | (~ spl30_49 | spl30_89)), inference(backward_demodulation, [], [f853, f592])).
fof(f853, plain, (~ (e3 = op(e0, op(e0, e3))) | spl30_89), inference(avatar_component_clause, [], [f852])).
fof(f852, plain, (spl30_89 <=> (e3 = op(e0, op(e0, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl30_89])])).
fof(f5517, plain, (spl30_49 | ~ spl30_64 | ~ spl30_92), inference(avatar_split_clause, [], [f5508, f867, f653, f590])).
fof(f5508, plain, ((e0 = op(e0, e3)) | (~ spl30_64 | ~ spl30_92)), inference(backward_demodulation, [], [f869, f655])).
fof(f869, plain, ((e0 = op(e0, op(e0, e0))) | ~ spl30_92), inference(avatar_component_clause, [], [f867])).
fof(f5516, plain, (~ spl30_52 | ~ spl30_64), inference(avatar_split_clause, [], [f5507, f653, f602])).
fof(f5507, plain, (~ (e3 = op(e0, e3)) | ~ spl30_64), inference(backward_demodulation, [], [f189, f655])).
fof(f5498, plain, (~ spl30_62 | ~ spl30_58), inference(avatar_split_clause, [], [f5497, f628, f645])).
fof(f5497, plain, (~ (op(e0, e0) = e1) | ~ spl30_58), inference(forward_demodulation, [], [f187, f630])).
fof(f5486, plain, (~ spl30_6 | ~ spl30_11 | spl30_68), inference(avatar_contradiction_clause, [], [f5485])).
fof(f5485, plain, ($false | (~ spl30_6 | ~ spl30_11 | spl30_68)), inference(subsumption_resolution, [], [f5482, f409])).
fof(f5482, plain, (~ (e1 = op(e3, e2)) | (~ spl30_11 | spl30_68)), inference(backward_demodulation, [], [f705, f430])).
fof(f705, plain, (~ (e1 = op(e3, op(e3, e1))) | spl30_68), inference(avatar_component_clause, [], [f704])).
fof(f704, plain, (spl30_68 <=> (e1 = op(e3, op(e3, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl30_68])])).
fof(f5479, plain, (spl30_1 | ~ spl30_16 | ~ spl30_104), inference(avatar_split_clause, [], [f5474, f943, f449, f386])).
fof(f943, plain, (spl30_104 <=> (e0 = op(op(e3, e0), e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_104])])).
fof(f5474, plain, ((e0 = op(e3, e3)) | (~ spl30_16 | ~ spl30_104)), inference(backward_demodulation, [], [f945, f451])).
fof(f945, plain, ((e0 = op(op(e3, e0), e3)) | ~ spl30_104), inference(avatar_component_clause, [], [f943])).
fof(f5478, plain, (~ spl30_16 | spl30_103), inference(avatar_contradiction_clause, [], [f5477])).
fof(f5477, plain, ($false | (~ spl30_16 | spl30_103)), inference(subsumption_resolution, [], [f5473, f451])).
fof(f5473, plain, (~ (e3 = op(e3, e0)) | (~ spl30_16 | spl30_103)), inference(backward_demodulation, [], [f940, f451])).
fof(f5475, plain, (~ spl30_4 | ~ spl30_16), inference(avatar_split_clause, [], [f5471, f449, f398])).
fof(f5471, plain, (~ (e3 = op(e3, e3)) | ~ spl30_16), inference(backward_demodulation, [], [f207, f451])).
fof(f5461, plain, (spl30_11 | ~ spl30_40 | ~ spl30_121), inference(avatar_split_clause, [], [f5458, f1029, f551, f428])).
fof(f5458, plain, ((e2 = op(e3, e1)) | (~ spl30_40 | ~ spl30_121)), inference(backward_demodulation, [], [f1031, f553])).
fof(f5452, plain, (~ spl30_33 | ~ spl30_41), inference(avatar_split_clause, [], [f5443, f556, f522])).
fof(f5443, plain, (~ (e0 = op(e1, e3)) | ~ spl30_41), inference(backward_demodulation, [], [f197, f558])).
fof(f5451, plain, (~ spl30_9 | ~ spl30_41), inference(avatar_split_clause, [], [f5441, f556, f420])).
fof(f5441, plain, (~ (e0 = op(e3, e1)) | ~ spl30_41), inference(backward_demodulation, [], [f173, f558])).
fof(f5440, plain, (~ spl30_39 | ~ spl30_47), inference(avatar_split_clause, [], [f5433, f581, f547])).
fof(f547, plain, (spl30_39 <=> (e2 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_39])])).
fof(f5433, plain, (~ (e2 = op(e1, e2)) | ~ spl30_47), inference(backward_demodulation, [], [f194, f583])).
fof(f5430, plain, (spl30_16 | ~ spl30_52 | ~ spl30_129), inference(avatar_split_clause, [], [f5427, f1070, f602, f449])).
fof(f5427, plain, ((e3 = op(e3, e0)) | (~ spl30_52 | ~ spl30_129)), inference(backward_demodulation, [], [f1072, f604])).
fof(f5419, plain, (~ spl30_39 | ~ spl30_55), inference(avatar_split_clause, [], [f5413, f615, f547])).
fof(f5413, plain, (~ (e2 = op(e1, e2)) | ~ spl30_55), inference(backward_demodulation, [], [f175, f617])).
fof(f5400, plain, (spl30_11 | ~ spl30_6 | ~ spl30_67), inference(avatar_split_clause, [], [f5399, f699, f407, f428])).
fof(f5399, plain, ((e2 = op(e3, e1)) | (~ spl30_6 | ~ spl30_67)), inference(forward_demodulation, [], [f701, f409])).
fof(f701, plain, ((e2 = op(e3, op(e3, e2))) | ~ spl30_67), inference(avatar_component_clause, [], [f699])).
fof(f5398, plain, (spl30_40 | ~ spl30_6 | ~ spl30_97), inference(avatar_split_clause, [], [f5397, f908, f407, f551])).
fof(f5397, plain, ((e3 = op(e1, e2)) | (~ spl30_6 | ~ spl30_97)), inference(forward_demodulation, [], [f909, f409])).
fof(f909, plain, ((e3 = op(op(e3, e2), e2)) | ~ spl30_97), inference(avatar_component_clause, [], [f908])).
fof(f5396, plain, (spl30_11 | ~ spl30_28 | ~ spl30_111), inference(avatar_split_clause, [], [f5395, f979, f500, f428])).
fof(f5395, plain, ((e2 = op(e3, e1)) | (~ spl30_28 | ~ spl30_111)), inference(forward_demodulation, [], [f980, f502])).
fof(f980, plain, ((e2 = op(op(e2, e1), e1)) | ~ spl30_111), inference(avatar_component_clause, [], [f979])).
fof(f5371, plain, (~ spl30_21 | ~ spl30_56 | spl30_147), inference(avatar_contradiction_clause, [], [f5370])).
fof(f5370, plain, ($false | (~ spl30_21 | ~ spl30_56 | spl30_147)), inference(subsumption_resolution, [], [f5369, f621])).
fof(f5369, plain, (~ (e3 = op(e0, e2)) | (~ spl30_21 | spl30_147)), inference(forward_demodulation, [], [f1169, f473])).
fof(f1169, plain, (~ (e3 = op(op(e2, e2), e2)) | spl30_147), inference(avatar_component_clause, [], [f1167])).
fof(f1167, plain, (spl30_147 <=> (e3 = op(op(e2, e2), e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_147])])).
fof(f5362, plain, (~ spl30_4 | spl30_66), inference(avatar_contradiction_clause, [], [f5361])).
fof(f5361, plain, ($false | (~ spl30_4 | spl30_66)), inference(subsumption_resolution, [], [f5355, f400])).
fof(f400, plain, ((e3 = op(e3, e3)) | ~ spl30_4), inference(avatar_component_clause, [], [f398])).
fof(f5355, plain, (~ (e3 = op(e3, e3)) | (~ spl30_4 | spl30_66)), inference(backward_demodulation, [], [f695, f400])).
fof(f695, plain, (~ (e3 = op(e3, op(e3, e3))) | spl30_66), inference(avatar_component_clause, [], [f694])).
fof(f694, plain, (spl30_66 <=> (e3 = op(e3, op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl30_66])])).
fof(f5345, plain, (~ spl30_3 | ~ spl30_15), inference(avatar_split_clause, [], [f5340, f445, f394])).
fof(f5340, plain, (~ (e2 = op(e3, e3)) | ~ spl30_15), inference(backward_demodulation, [], [f207, f447])).
fof(f5330, plain, (~ spl30_6 | ~ spl30_21 | ~ spl30_56 | spl30_148), inference(avatar_contradiction_clause, [], [f5329])).
fof(f5329, plain, ($false | (~ spl30_6 | ~ spl30_21 | ~ spl30_56 | spl30_148)), inference(subsumption_resolution, [], [f5328, f409])).
fof(f5328, plain, (~ (e1 = op(e3, e2)) | (~ spl30_21 | ~ spl30_56 | spl30_148)), inference(forward_demodulation, [], [f5325, f621])).
fof(f5325, plain, (~ (e1 = op(op(e0, e2), e2)) | (~ spl30_21 | spl30_148)), inference(backward_demodulation, [], [f1174, f473])).
fof(f1174, plain, (~ (e1 = op(op(op(e2, e2), e2), e2)) | spl30_148), inference(avatar_component_clause, [], [f1172])).
fof(f1172, plain, (spl30_148 <=> (e1 = op(op(op(e2, e2), e2), e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_148])])).
fof(f5315, plain, (~ spl30_12 | ~ spl30_28), inference(avatar_split_clause, [], [f5310, f500, f432])).
fof(f5310, plain, (~ (e3 = op(e3, e1)) | ~ spl30_28), inference(backward_demodulation, [], [f174, f502])).
fof(f5309, plain, (~ spl30_23 | ~ spl30_39), inference(avatar_split_clause, [], [f5304, f547, f479])).
fof(f5304, plain, (~ (e2 = op(e2, e2)) | ~ spl30_39), inference(backward_demodulation, [], [f178, f549])).
fof(f549, plain, ((e2 = op(e1, e2)) | ~ spl30_39), inference(avatar_component_clause, [], [f547])).
fof(f5283, plain, (~ spl30_61 | spl30_92), inference(avatar_contradiction_clause, [], [f5282])).
fof(f5282, plain, ($false | (~ spl30_61 | spl30_92)), inference(subsumption_resolution, [], [f5274, f643])).
fof(f5274, plain, (~ (e0 = op(e0, e0)) | (~ spl30_61 | spl30_92)), inference(backward_demodulation, [], [f868, f643])).
fof(f5281, plain, (~ spl30_53 | ~ spl30_61), inference(avatar_split_clause, [], [f5270, f641, f607])).
fof(f607, plain, (spl30_53 <=> (e0 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_53])])).
fof(f5270, plain, (~ (e0 = op(e0, e2)) | ~ spl30_61), inference(backward_demodulation, [], [f188, f643])).
fof(f5266, plain, (~ spl30_25 | ~ spl30_30 | spl30_76), inference(avatar_split_clause, [], [f5265, f755, f509, f488])).
fof(f5265, plain, (~ (e0 = op(e2, e1)) | (~ spl30_30 | spl30_76)), inference(forward_demodulation, [], [f756, f511])).
fof(f5264, plain, (spl30_48 | ~ spl30_33 | ~ spl30_81), inference(avatar_split_clause, [], [f5263, f796, f522, f585])).
fof(f5263, plain, ((e3 = op(e1, e0)) | (~ spl30_33 | ~ spl30_81)), inference(forward_demodulation, [], [f798, f524])).
fof(f798, plain, ((e3 = op(e1, op(e1, e3))) | ~ spl30_81), inference(avatar_component_clause, [], [f796])).
fof(f5238, plain, (~ spl30_4 | ~ spl30_12), inference(avatar_split_clause, [], [f5234, f432, f398])).
fof(f5234, plain, (~ (e3 = op(e3, e3)) | ~ spl30_12), inference(backward_demodulation, [], [f209, f434])).
fof(f5231, plain, (~ spl30_4 | ~ spl30_20), inference(avatar_split_clause, [], [f5227, f466, f398])).
fof(f5227, plain, (~ (e3 = op(e3, e3)) | ~ spl30_20), inference(backward_demodulation, [], [f186, f468])).
fof(f468, plain, ((e3 = op(e2, e3)) | ~ spl30_20), inference(avatar_component_clause, [], [f466])).
fof(f5217, plain, (spl30_25 | ~ spl30_30 | ~ spl30_76), inference(avatar_split_clause, [], [f5213, f755, f509, f488])).
fof(f5213, plain, ((e0 = op(e2, e1)) | (~ spl30_30 | ~ spl30_76)), inference(backward_demodulation, [], [f757, f511])).
fof(f5216, plain, (~ spl30_26 | ~ spl30_30), inference(avatar_split_clause, [], [f5209, f509, f492])).
fof(f5209, plain, (~ (e1 = op(e2, e1)) | ~ spl30_30), inference(backward_demodulation, [], [f199, f511])).
fof(f5206, plain, (~ spl30_17 | ~ spl30_33), inference(avatar_split_clause, [], [f5198, f522, f454])).
fof(f5198, plain, (~ (e0 = op(e2, e3)) | ~ spl30_33), inference(backward_demodulation, [], [f184, f524])).
fof(f5169, plain, (spl30_30 | ~ spl30_47 | ~ spl30_125), inference(avatar_split_clause, [], [f5165, f1050, f581, f509])).
fof(f5165, plain, ((e1 = op(e2, e0)) | (~ spl30_47 | ~ spl30_125)), inference(backward_demodulation, [], [f1051, f583])).
fof(f5158, plain, (~ spl30_37 | ~ spl30_53), inference(avatar_split_clause, [], [f5152, f607, f539])).
fof(f5152, plain, (~ (e0 = op(e1, e2)) | ~ spl30_53), inference(backward_demodulation, [], [f175, f609])).
fof(f609, plain, ((e0 = op(e0, e2)) | ~ spl30_53), inference(avatar_component_clause, [], [f607])).
fof(f5151, plain, (spl30_30 | ~ spl30_59 | ~ spl30_135), inference(avatar_split_clause, [], [f5148, f1100, f632, f509])).
fof(f5148, plain, ((e1 = op(e2, e0)) | (~ spl30_59 | ~ spl30_135)), inference(backward_demodulation, [], [f1102, f634])).
fof(f5140, plain, (~ spl30_13 | ~ spl30_64 | spl30_162), inference(avatar_contradiction_clause, [], [f5139])).
fof(f5139, plain, ($false | (~ spl30_13 | ~ spl30_64 | spl30_162)), inference(subsumption_resolution, [], [f5138, f655])).
fof(f5138, plain, (~ (op(e0, e0) = e3) | (~ spl30_13 | ~ spl30_64 | spl30_162)), inference(forward_demodulation, [], [f5131, f439])).
fof(f5131, plain, (~ (e3 = op(op(e3, e0), e0)) | (~ spl30_64 | spl30_162)), inference(backward_demodulation, [], [f1244, f655])).
fof(f1244, plain, (~ (e3 = op(op(op(e0, e0), e0), e0)) | spl30_162), inference(avatar_component_clause, [], [f1242])).
fof(f5133, plain, (~ spl30_56 | ~ spl30_64), inference(avatar_split_clause, [], [f5125, f653, f619])).
fof(f5125, plain, (~ (e3 = op(e0, e2)) | ~ spl30_64), inference(backward_demodulation, [], [f188, f655])).
fof(f5132, plain, (~ spl30_32 | ~ spl30_64), inference(avatar_split_clause, [], [f5123, f653, f517])).
fof(f5123, plain, (~ (e3 = op(e2, e0)) | ~ spl30_64), inference(backward_demodulation, [], [f164, f655])).
fof(f5117, plain, (spl30_33 | ~ spl30_50 | ~ spl30_128), inference(avatar_split_clause, [], [f5116, f1065, f594, f522])).
fof(f5116, plain, ((e0 = op(e1, e3)) | (~ spl30_50 | ~ spl30_128)), inference(forward_demodulation, [], [f1066, f596])).
fof(f1066, plain, ((e0 = op(op(e0, e3), e3)) | ~ spl30_128), inference(avatar_component_clause, [], [f1065])).
fof(f5109, plain, (~ spl30_60 | ~ spl30_50 | spl30_89), inference(avatar_split_clause, [], [f5024, f852, f594, f636])).
fof(f5024, plain, (~ (e3 = op(e0, e1)) | (~ spl30_50 | spl30_89)), inference(backward_demodulation, [], [f853, f596])).
fof(f5053, plain, (~ spl30_22 | ~ spl30_26), inference(avatar_split_clause, [], [f5049, f492, f475])).
fof(f5049, plain, (~ (e1 = op(e2, e2)) | ~ spl30_26), inference(backward_demodulation, [], [f202, f494])).
fof(f5039, plain, (spl30_41 | ~ spl30_46 | ~ spl30_84), inference(avatar_contradiction_clause, [], [f5038])).
fof(f5038, plain, ($false | (spl30_41 | ~ spl30_46 | ~ spl30_84)), inference(subsumption_resolution, [], [f5033, f557])).
fof(f557, plain, (~ (e0 = op(e1, e1)) | spl30_41), inference(avatar_component_clause, [], [f556])).
fof(f5033, plain, ((e0 = op(e1, e1)) | (~ spl30_46 | ~ spl30_84)), inference(backward_demodulation, [], [f813, f579])).
fof(f5037, plain, (~ spl30_34 | ~ spl30_46), inference(avatar_split_clause, [], [f5030, f577, f526])).
fof(f5030, plain, (~ (e1 = op(e1, e3)) | ~ spl30_46), inference(backward_demodulation, [], [f195, f579])).
fof(f5036, plain, (~ spl30_14 | ~ spl30_46), inference(avatar_split_clause, [], [f5029, f577, f441])).
fof(f5029, plain, (~ (e1 = op(e3, e0)) | ~ spl30_46), inference(backward_demodulation, [], [f167, f579])).
fof(f5018, plain, (spl30_15 | ~ spl30_56 | ~ spl30_132), inference(avatar_contradiction_clause, [], [f5017])).
fof(f5017, plain, ($false | (spl30_15 | ~ spl30_56 | ~ spl30_132)), inference(subsumption_resolution, [], [f5014, f446])).
fof(f446, plain, (~ (e2 = op(e3, e0)) | spl30_15), inference(avatar_component_clause, [], [f445])).
fof(f5014, plain, ((e2 = op(e3, e0)) | (~ spl30_56 | ~ spl30_132)), inference(backward_demodulation, [], [f1087, f621])).
fof(f5016, plain, (~ spl30_52 | ~ spl30_56), inference(avatar_split_clause, [], [f5011, f619, f602])).
fof(f5011, plain, (~ (e3 = op(e0, e3)) | ~ spl30_56), inference(backward_demodulation, [], [f192, f621])).
fof(f5015, plain, (~ spl30_8 | ~ spl30_56), inference(avatar_split_clause, [], [f5010, f619, f415])).
fof(f5010, plain, (~ (e3 = op(e3, e2)) | ~ spl30_56), inference(backward_demodulation, [], [f177, f621])).
fof(f4996, plain, (~ spl30_32 | ~ spl30_63 | spl30_154), inference(avatar_contradiction_clause, [], [f4995])).
fof(f4995, plain, ($false | (~ spl30_32 | ~ spl30_63 | spl30_154)), inference(subsumption_resolution, [], [f4988, f519])).
fof(f4988, plain, (~ (e3 = op(e2, e0)) | (~ spl30_63 | spl30_154)), inference(backward_demodulation, [], [f1202, f651])).
fof(f1202, plain, (~ (e3 = op(op(e0, e0), e0)) | spl30_154), inference(avatar_component_clause, [], [f1200])).
fof(f1200, plain, (spl30_154 <=> (e3 = op(op(e0, e0), e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_154])])).
fof(f4994, plain, (~ spl30_14 | ~ spl30_32 | ~ spl30_63 | spl30_152), inference(avatar_split_clause, [], [f4993, f1190, f649, f517, f441])).
fof(f4993, plain, (~ (e1 = op(e3, e0)) | (~ spl30_32 | ~ spl30_63 | spl30_152)), inference(forward_demodulation, [], [f4986, f519])).
fof(f4986, plain, (~ (e1 = op(op(e2, e0), e0)) | (~ spl30_63 | spl30_152)), inference(backward_demodulation, [], [f1192, f651])).
fof(f4978, plain, (spl30_6 | ~ spl30_11 | ~ spl30_68), inference(avatar_split_clause, [], [f4977, f704, f428, f407])).
fof(f4977, plain, ((e1 = op(e3, e2)) | (~ spl30_11 | ~ spl30_68)), inference(forward_demodulation, [], [f706, f430])).
fof(f706, plain, ((e1 = op(e3, op(e3, e1))) | ~ spl30_68), inference(avatar_component_clause, [], [f704])).
fof(f4974, plain, (~ spl30_34 | ~ spl30_44 | spl30_83), inference(avatar_split_clause, [], [f4973, f806, f568, f526])).
fof(f806, plain, (spl30_83 <=> (e1 = op(e1, op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl30_83])])).
fof(f4973, plain, (~ (e1 = op(e1, e3)) | (~ spl30_44 | spl30_83)), inference(forward_demodulation, [], [f807, f570])).
fof(f570, plain, ((e3 = op(e1, e1)) | ~ spl30_44), inference(avatar_component_clause, [], [f568])).
fof(f807, plain, (~ (e1 = op(e1, op(e1, e1))) | spl30_83), inference(avatar_component_clause, [], [f806])).
fof(f4963, plain, (~ spl30_36 | ~ spl30_44), inference(avatar_split_clause, [], [f4962, f568, f534])).
fof(f4962, plain, (~ (e3 = op(e1, e3)) | ~ spl30_44), inference(forward_demodulation, [], [f197, f570])).
fof(f4958, plain, (~ spl30_16 | ~ spl30_32), inference(avatar_split_clause, [], [f4957, f517, f449])).
fof(f4957, plain, (~ (e3 = op(e3, e0)) | ~ spl30_32), inference(forward_demodulation, [], [f168, f519])).
fof(f4942, plain, (~ spl30_11 | ~ spl30_25 | ~ spl30_100), inference(avatar_contradiction_clause, [], [f4941])).
fof(f4941, plain, ($false | (~ spl30_11 | ~ spl30_25 | ~ spl30_100)), inference(subsumption_resolution, [], [f4940, f213])).
fof(f213, plain, ~ (e0 = e3), inference(cnf_transformation, [], [f4])).
fof(f4940, plain, ((e0 = e3) | (~ spl30_11 | ~ spl30_25 | ~ spl30_100)), inference(forward_demodulation, [], [f4937, f490])).
fof(f4937, plain, ((e3 = op(e2, e1)) | (~ spl30_11 | ~ spl30_100)), inference(backward_demodulation, [], [f924, f430])).
fof(f924, plain, ((e3 = op(op(e3, e1), e1)) | ~ spl30_100), inference(avatar_component_clause, [], [f923])).
fof(f923, plain, (spl30_100 <=> (e3 = op(op(e3, e1), e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_100])])).
fof(f4918, plain, (~ spl30_25 | spl30_54 | ~ spl30_112), inference(avatar_contradiction_clause, [], [f4917])).
fof(f4917, plain, ($false | (~ spl30_25 | spl30_54 | ~ spl30_112)), inference(subsumption_resolution, [], [f4914, f612])).
fof(f612, plain, (~ (e1 = op(e0, e2)) | spl30_54), inference(avatar_component_clause, [], [f611])).
fof(f4914, plain, ((e1 = op(e0, e2)) | (~ spl30_25 | ~ spl30_112)), inference(backward_demodulation, [], [f986, f490])).
fof(f986, plain, ((e1 = op(op(e2, e1), e2)) | ~ spl30_112), inference(avatar_component_clause, [], [f984])).
fof(f4904, plain, (~ spl30_37 | ~ spl30_47 | spl30_82), inference(avatar_contradiction_clause, [], [f4903])).
fof(f4903, plain, ($false | (~ spl30_37 | ~ spl30_47 | spl30_82)), inference(subsumption_resolution, [], [f4899, f583])).
fof(f4899, plain, (~ (e2 = op(e1, e0)) | (~ spl30_37 | spl30_82)), inference(backward_demodulation, [], [f802, f541])).
fof(f4893, plain, (spl30_11 | ~ spl30_44 | ~ spl30_144), inference(avatar_split_clause, [], [f4889, f1152, f568, f428])).
fof(f4889, plain, ((e2 = op(e3, e1)) | (~ spl30_44 | ~ spl30_144)), inference(backward_demodulation, [], [f1153, f570])).
fof(f4892, plain, (~ spl30_40 | ~ spl30_44), inference(avatar_split_clause, [], [f4887, f568, f551])).
fof(f4887, plain, (~ (e3 = op(e1, e2)) | ~ spl30_44), inference(backward_demodulation, [], [f196, f570])).
fof(f4877, plain, (~ spl30_41 | ~ spl30_58 | spl30_134), inference(avatar_split_clause, [], [f4873, f1095, f628, f556])).
fof(f4873, plain, (~ (e0 = op(e1, e1)) | (~ spl30_58 | spl30_134)), inference(backward_demodulation, [], [f1097, f630])).
fof(f4876, plain, (~ spl30_54 | ~ spl30_58), inference(avatar_split_clause, [], [f4869, f628, f611])).
fof(f4869, plain, (~ (e1 = op(e0, e2)) | ~ spl30_58), inference(backward_demodulation, [], [f190, f630])).
fof(f190, plain, ~ (op(e0, e1) = op(e0, e2)), inference(cnf_transformation, [], [f3])).
fof(f4875, plain, (~ spl30_26 | ~ spl30_58), inference(avatar_split_clause, [], [f4867, f628, f492])).
fof(f4867, plain, (~ (e1 = op(e2, e1)) | ~ spl30_58), inference(backward_demodulation, [], [f170, f630])).
fof(f4840, plain, (spl30_44 | ~ spl30_34 | ~ spl30_81), inference(avatar_split_clause, [], [f4764, f796, f526, f568])).
fof(f4764, plain, ((e3 = op(e1, e1)) | (~ spl30_34 | ~ spl30_81)), inference(backward_demodulation, [], [f798, f528])).
fof(f4835, plain, (~ spl30_20 | ~ spl30_32), inference(avatar_split_clause, [], [f4834, f517, f466])).
fof(f4834, plain, (~ (e3 = op(e2, e3)) | ~ spl30_32), inference(forward_demodulation, [], [f201, f519])).
fof(f4828, plain, (~ spl30_3 | ~ spl30_17 | spl30_151), inference(avatar_contradiction_clause, [], [f4827])).
fof(f4827, plain, ($false | (~ spl30_3 | ~ spl30_17 | spl30_151)), inference(subsumption_resolution, [], [f4820, f456])).
fof(f4820, plain, (~ (e0 = op(e2, e3)) | (~ spl30_3 | spl30_151)), inference(backward_demodulation, [], [f1187, f396])).
fof(f396, plain, ((e2 = op(e3, e3)) | ~ spl30_3), inference(avatar_component_clause, [], [f394])).
fof(f1187, plain, (~ (e0 = op(op(e3, e3), e3)) | spl30_151), inference(avatar_component_clause, [], [f1185])).
fof(f1185, plain, (spl30_151 <=> (e0 = op(op(e3, e3), e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_151])])).
fof(f4824, plain, (~ spl30_3 | ~ spl30_5 | ~ spl30_66), inference(avatar_contradiction_clause, [], [f4823])).
fof(f4823, plain, ($false | (~ spl30_3 | ~ spl30_5 | ~ spl30_66)), inference(subsumption_resolution, [], [f4822, f213])).
fof(f4822, plain, ((e0 = e3) | (~ spl30_3 | ~ spl30_5 | ~ spl30_66)), inference(forward_demodulation, [], [f4817, f405])).
fof(f4817, plain, ((e3 = op(e3, e2)) | (~ spl30_3 | ~ spl30_66)), inference(backward_demodulation, [], [f696, f396])).
fof(f696, plain, ((e3 = op(e3, op(e3, e3))) | ~ spl30_66), inference(avatar_component_clause, [], [f694])).
fof(f4806, plain, (spl30_2 | ~ spl30_12 | ~ spl30_68), inference(avatar_contradiction_clause, [], [f4805])).
fof(f4805, plain, ($false | (spl30_2 | ~ spl30_12 | ~ spl30_68)), inference(subsumption_resolution, [], [f4801, f391])).
fof(f391, plain, (~ (e1 = op(e3, e3)) | spl30_2), inference(avatar_component_clause, [], [f390])).
fof(f4801, plain, ((e1 = op(e3, e3)) | (~ spl30_12 | ~ spl30_68)), inference(backward_demodulation, [], [f706, f434])).
fof(f4795, plain, (~ spl30_23 | spl30_74), inference(avatar_contradiction_clause, [], [f4794])).
fof(f4794, plain, ($false | (~ spl30_23 | spl30_74)), inference(subsumption_resolution, [], [f4790, f481])).
fof(f4790, plain, (~ (e2 = op(e2, e2)) | (~ spl30_23 | spl30_74)), inference(backward_demodulation, [], [f746, f481])).
fof(f746, plain, (~ (e2 = op(e2, op(e2, e2))) | spl30_74), inference(avatar_component_clause, [], [f745])).
fof(f4760, plain, (~ spl30_28 | ~ spl30_41 | ~ spl30_59 | spl30_160), inference(avatar_split_clause, [], [f4759, f1232, f632, f556, f500])).
fof(f1232, plain, (spl30_160 <=> (e3 = op(op(op(e1, e1), e1), e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_160])])).
fof(f4759, plain, (~ (e3 = op(e2, e1)) | (~ spl30_41 | ~ spl30_59 | spl30_160)), inference(forward_demodulation, [], [f4751, f634])).
fof(f4751, plain, (~ (e3 = op(op(e0, e1), e1)) | (~ spl30_41 | spl30_160)), inference(backward_demodulation, [], [f1234, f558])).
fof(f1234, plain, (~ (e3 = op(op(op(e1, e1), e1), e1)) | spl30_160), inference(avatar_component_clause, [], [f1232])).
fof(f4757, plain, (~ spl30_41 | ~ spl30_59 | spl30_144), inference(avatar_contradiction_clause, [], [f4756])).
fof(f4756, plain, ($false | (~ spl30_41 | ~ spl30_59 | spl30_144)), inference(subsumption_resolution, [], [f4748, f634])).
fof(f4748, plain, (~ (e2 = op(e0, e1)) | (~ spl30_41 | spl30_144)), inference(backward_demodulation, [], [f1154, f558])).
fof(f1154, plain, (~ (e2 = op(op(e1, e1), e1)) | spl30_144), inference(avatar_component_clause, [], [f1152])).
fof(f4754, plain, (~ spl30_41 | spl30_46 | ~ spl30_83), inference(avatar_contradiction_clause, [], [f4753])).
fof(f4753, plain, ($false | (~ spl30_41 | spl30_46 | ~ spl30_83)), inference(subsumption_resolution, [], [f4746, f578])).
fof(f578, plain, (~ (e1 = op(e1, e0)) | spl30_46), inference(avatar_component_clause, [], [f577])).
fof(f4746, plain, ((e1 = op(e1, e0)) | (~ spl30_41 | ~ spl30_83)), inference(backward_demodulation, [], [f808, f558])).
fof(f808, plain, ((e1 = op(e1, op(e1, e1))) | ~ spl30_83), inference(avatar_component_clause, [], [f806])).
fof(f4707, plain, (~ spl30_45 | ~ spl30_61), inference(avatar_split_clause, [], [f4694, f641, f573])).
fof(f4694, plain, (~ (e0 = op(e1, e0)) | ~ spl30_61), inference(backward_demodulation, [], [f163, f643])).
fof(f4691, plain, (spl30_32 | ~ spl30_17 | ~ spl30_73), inference(avatar_split_clause, [], [f4690, f740, f454, f517])).
fof(f4690, plain, ((e3 = op(e2, e0)) | (~ spl30_17 | ~ spl30_73)), inference(forward_demodulation, [], [f742, f456])).
fof(f742, plain, ((e3 = op(e2, op(e2, e3))) | ~ spl30_73), inference(avatar_component_clause, [], [f740])).
fof(f4659, plain, (~ spl30_45 | spl30_57 | ~ spl30_126), inference(avatar_contradiction_clause, [], [f4658])).
fof(f4658, plain, ($false | (~ spl30_45 | spl30_57 | ~ spl30_126)), inference(subsumption_resolution, [], [f4654, f625])).
fof(f625, plain, (~ (e0 = op(e0, e1)) | spl30_57), inference(avatar_component_clause, [], [f624])).
fof(f4654, plain, ((e0 = op(e0, e1)) | (~ spl30_45 | ~ spl30_126)), inference(backward_demodulation, [], [f1057, f575])).
fof(f4657, plain, (~ spl30_45 | spl30_62 | ~ spl30_125), inference(avatar_contradiction_clause, [], [f4656])).
fof(f4656, plain, ($false | (~ spl30_45 | spl30_62 | ~ spl30_125)), inference(subsumption_resolution, [], [f4653, f646])).
fof(f646, plain, (~ (op(e0, e0) = e1) | spl30_62), inference(avatar_component_clause, [], [f645])).
fof(f4653, plain, ((op(e0, e0) = e1) | (~ spl30_45 | ~ spl30_125)), inference(backward_demodulation, [], [f1051, f575])).
fof(f4639, plain, (~ spl30_48 | ~ spl30_64), inference(avatar_split_clause, [], [f4635, f653, f585])).
fof(f4635, plain, (~ (e3 = op(e1, e0)) | ~ spl30_64), inference(backward_demodulation, [], [f163, f655])).
fof(f4634, plain, (spl30_40 | ~ spl30_35 | ~ spl30_81), inference(avatar_split_clause, [], [f4633, f796, f530, f551])).
fof(f4633, plain, ((e3 = op(e1, e2)) | (~ spl30_35 | ~ spl30_81)), inference(forward_demodulation, [], [f798, f532])).
fof(f4618, plain, (~ spl30_48 | ~ spl30_159 | spl30_162), inference(avatar_split_clause, [], [f4617, f1242, f1224, f585])).
fof(f4617, plain, (~ (e3 = op(e1, e0)) | (~ spl30_159 | spl30_162)), inference(backward_demodulation, [], [f1244, f1225])).
fof(f4604, plain, (spl30_56 | ~ spl30_17 | ~ spl30_107), inference(avatar_split_clause, [], [f4564, f958, f454, f619])).
fof(f958, plain, (spl30_107 <=> (e3 = op(op(e2, e3), e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_107])])).
fof(f4564, plain, ((e3 = op(e0, e2)) | (~ spl30_17 | ~ spl30_107)), inference(backward_demodulation, [], [f960, f456])).
fof(f960, plain, ((e3 = op(op(e2, e3), e2)) | ~ spl30_107), inference(avatar_component_clause, [], [f958])).
fof(f4598, plain, (~ spl30_40 | ~ spl30_22 | spl30_147), inference(avatar_split_clause, [], [f4547, f1167, f475, f551])).
fof(f4547, plain, (~ (e3 = op(e1, e2)) | (~ spl30_22 | spl30_147)), inference(backward_demodulation, [], [f1169, f477])).
fof(f477, plain, ((e1 = op(e2, e2)) | ~ spl30_22), inference(avatar_component_clause, [], [f475])).
fof(f4596, plain, (~ spl30_3 | ~ spl30_4), inference(avatar_contradiction_clause, [], [f4595])).
fof(f4595, plain, ($false | (~ spl30_3 | ~ spl30_4)), inference(subsumption_resolution, [], [f4594, f216])).
fof(f216, plain, ~ (e2 = e3), inference(cnf_transformation, [], [f4])).
fof(f4594, plain, ((e2 = e3) | (~ spl30_3 | ~ spl30_4)), inference(backward_demodulation, [], [f400, f396])).
fof(f4583, plain, (~ spl30_3 | ~ spl30_7), inference(avatar_split_clause, [], [f4578, f411, f394])).
fof(f4578, plain, (~ (e2 = op(e3, e3)) | ~ spl30_7), inference(backward_demodulation, [], [f210, f413])).
fof(f413, plain, ((e2 = op(e3, e2)) | ~ spl30_7), inference(avatar_component_clause, [], [f411])).
fof(f4576, plain, (~ spl30_1 | ~ spl30_9), inference(avatar_split_clause, [], [f4569, f420, f386])).
fof(f4569, plain, (~ (e0 = op(e3, e3)) | ~ spl30_9), inference(backward_demodulation, [], [f209, f422])).
fof(f4557, plain, (~ spl30_22 | ~ spl30_37 | ~ spl30_56 | spl30_161), inference(avatar_contradiction_clause, [], [f4556])).
fof(f4556, plain, ($false | (~ spl30_22 | ~ spl30_37 | ~ spl30_56 | spl30_161)), inference(subsumption_resolution, [], [f4555, f621])).
fof(f4555, plain, (~ (e3 = op(e0, e2)) | (~ spl30_22 | ~ spl30_37 | spl30_161)), inference(forward_demodulation, [], [f4550, f541])).
fof(f4550, plain, (~ (e3 = op(op(e1, e2), e2)) | (~ spl30_22 | spl30_161)), inference(backward_demodulation, [], [f1239, f477])).
fof(f4554, plain, (~ spl30_22 | ~ spl30_37 | spl30_149), inference(avatar_contradiction_clause, [], [f4553])).
fof(f4553, plain, ($false | (~ spl30_22 | ~ spl30_37 | spl30_149)), inference(subsumption_resolution, [], [f4549, f541])).
fof(f4549, plain, (~ (e0 = op(e1, e2)) | (~ spl30_22 | spl30_149)), inference(backward_demodulation, [], [f1178, f477])).
fof(f1178, plain, (~ (e0 = op(op(e2, e2), e2)) | spl30_149), inference(avatar_component_clause, [], [f1176])).
fof(f1176, plain, (spl30_149 <=> (e0 = op(op(e2, e2), e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_149])])).
fof(f4541, plain, (spl30_6 | ~ spl30_28 | ~ spl30_112), inference(avatar_contradiction_clause, [], [f4540])).
fof(f4540, plain, ($false | (spl30_6 | ~ spl30_28 | ~ spl30_112)), inference(subsumption_resolution, [], [f4538, f408])).
fof(f4538, plain, ((e1 = op(e3, e2)) | (~ spl30_28 | ~ spl30_112)), inference(backward_demodulation, [], [f986, f502])).
fof(f4533, plain, (spl30_21 | ~ spl30_31 | ~ spl30_115), inference(avatar_contradiction_clause, [], [f4532])).
fof(f4532, plain, ($false | (spl30_21 | ~ spl30_31 | ~ spl30_115)), inference(subsumption_resolution, [], [f4530, f472])).
fof(f472, plain, (~ (e0 = op(e2, e2)) | spl30_21), inference(avatar_component_clause, [], [f471])).
fof(f4530, plain, ((e0 = op(e2, e2)) | (~ spl30_31 | ~ spl30_115)), inference(backward_demodulation, [], [f1001, f515])).
fof(f4508, plain, (~ spl30_38 | ~ spl30_42), inference(avatar_split_clause, [], [f4498, f560, f543])).
fof(f4498, plain, (~ (e1 = op(e1, e2)) | ~ spl30_42), inference(backward_demodulation, [], [f196, f562])).
fof(f4495, plain, (~ spl30_14 | ~ spl30_56 | ~ spl30_132), inference(avatar_contradiction_clause, [], [f4494])).
fof(f4494, plain, ($false | (~ spl30_14 | ~ spl30_56 | ~ spl30_132)), inference(subsumption_resolution, [], [f4493, f214])).
fof(f214, plain, ~ (e1 = e2), inference(cnf_transformation, [], [f4])).
fof(f4493, plain, ((e1 = e2) | (~ spl30_14 | ~ spl30_56 | ~ spl30_132)), inference(forward_demodulation, [], [f4490, f443])).
fof(f4490, plain, ((e2 = op(e3, e0)) | (~ spl30_56 | ~ spl30_132)), inference(backward_demodulation, [], [f1087, f621])).
fof(f4438, plain, (~ spl30_39 | ~ spl30_35), inference(avatar_split_clause, [], [f4437, f530, f547])).
fof(f4437, plain, (~ (e2 = op(e1, e2)) | ~ spl30_35), inference(forward_demodulation, [], [f198, f532])).
fof(f4431, plain, (spl30_28 | ~ spl30_35 | ~ spl30_118), inference(avatar_split_clause, [], [f4413, f1014, f530, f500])).
fof(f4413, plain, ((e3 = op(e2, e1)) | (~ spl30_35 | ~ spl30_118)), inference(forward_demodulation, [], [f1016, f532])).
fof(f4430, plain, (~ spl30_3 | ~ spl30_35), inference(avatar_split_clause, [], [f4320, f530, f394])).
fof(f4320, plain, (~ (e2 = op(e3, e3)) | ~ spl30_35), inference(backward_demodulation, [], [f185, f532])).
fof(f185, plain, ~ (op(e1, e3) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f4416, plain, (~ spl30_26 | ~ spl30_35 | ~ spl30_118), inference(avatar_contradiction_clause, [], [f4415])).
fof(f4415, plain, ($false | (~ spl30_26 | ~ spl30_35 | ~ spl30_118)), inference(subsumption_resolution, [], [f4414, f215])).
fof(f4414, plain, ((e1 = e3) | (~ spl30_26 | ~ spl30_35 | ~ spl30_118)), inference(forward_demodulation, [], [f4413, f494])).
fof(f4409, plain, (~ spl30_1 | ~ spl30_35 | ~ spl30_50 | spl30_157), inference(avatar_contradiction_clause, [], [f4408])).
fof(f4408, plain, ($false | (~ spl30_1 | ~ spl30_35 | ~ spl30_50 | spl30_157)), inference(subsumption_resolution, [], [f4407, f532])).
fof(f4407, plain, (~ (e2 = op(e1, e3)) | (~ spl30_1 | ~ spl30_50 | spl30_157)), inference(forward_demodulation, [], [f4406, f596])).
fof(f4406, plain, (~ (e2 = op(op(e0, e3), e3)) | (~ spl30_1 | spl30_157)), inference(forward_demodulation, [], [f1217, f388])).
fof(f1217, plain, (~ (e2 = op(op(op(e3, e3), e3), e3)) | spl30_157), inference(avatar_component_clause, [], [f1215])).
fof(f1215, plain, (spl30_157 <=> (e2 = op(op(op(e3, e3), e3), e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_157])])).
fof(f4384, plain, (~ spl30_41 | ~ spl30_60 | spl30_146), inference(avatar_contradiction_clause, [], [f4383])).
fof(f4383, plain, ($false | (~ spl30_41 | ~ spl30_60 | spl30_146)), inference(subsumption_resolution, [], [f4382, f638])).
fof(f4382, plain, (~ (e3 = op(e0, e1)) | (~ spl30_41 | spl30_146)), inference(forward_demodulation, [], [f1164, f558])).
fof(f4380, plain, (~ spl30_1 | ~ spl30_50 | spl30_142), inference(avatar_contradiction_clause, [], [f4379])).
fof(f4379, plain, ($false | (~ spl30_1 | ~ spl30_50 | spl30_142)), inference(subsumption_resolution, [], [f4372, f596])).
fof(f4372, plain, (~ (e1 = op(e0, e3)) | (~ spl30_1 | spl30_142)), inference(backward_demodulation, [], [f1145, f388])).
fof(f1145, plain, (~ (e1 = op(op(e3, e3), e3)) | spl30_142), inference(avatar_component_clause, [], [f1143])).
fof(f1143, plain, (spl30_142 <=> (e1 = op(op(e3, e3), e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_142])])).
fof(f4377, plain, (~ spl30_1 | ~ spl30_14 | ~ spl30_66), inference(avatar_contradiction_clause, [], [f4376])).
fof(f4376, plain, ($false | (~ spl30_1 | ~ spl30_14 | ~ spl30_66)), inference(subsumption_resolution, [], [f4375, f215])).
fof(f4375, plain, ((e1 = e3) | (~ spl30_1 | ~ spl30_14 | ~ spl30_66)), inference(forward_demodulation, [], [f4370, f443])).
fof(f4370, plain, ((e3 = op(e3, e0)) | (~ spl30_1 | ~ spl30_66)), inference(backward_demodulation, [], [f696, f388])).
fof(f4366, plain, (spl30_3 | ~ spl30_8 | ~ spl30_67), inference(avatar_contradiction_clause, [], [f4365])).
fof(f4365, plain, ($false | (spl30_3 | ~ spl30_8 | ~ spl30_67)), inference(subsumption_resolution, [], [f4361, f395])).
fof(f395, plain, (~ (e2 = op(e3, e3)) | spl30_3), inference(avatar_component_clause, [], [f394])).
fof(f4361, plain, ((e2 = op(e3, e3)) | (~ spl30_8 | ~ spl30_67)), inference(backward_demodulation, [], [f701, f417])).
fof(f4356, plain, (~ spl30_7 | ~ spl30_11), inference(avatar_split_clause, [], [f4350, f428, f411])).
fof(f4350, plain, (~ (e2 = op(e3, e2)) | ~ spl30_11), inference(backward_demodulation, [], [f208, f430])).
fof(f4306, plain, (~ spl30_33 | ~ spl30_48 | spl30_84), inference(avatar_split_clause, [], [f4302, f811, f585, f522])).
fof(f4302, plain, (~ (e0 = op(e1, e3)) | (~ spl30_48 | spl30_84)), inference(backward_demodulation, [], [f812, f587])).
fof(f4305, plain, (~ spl30_44 | ~ spl30_48), inference(avatar_split_clause, [], [f4300, f585, f568])).
fof(f4300, plain, (~ (e3 = op(e1, e1)) | ~ spl30_48), inference(backward_demodulation, [], [f193, f587])).
fof(f4297, plain, (~ spl30_9 | ~ spl30_60 | spl30_134), inference(avatar_split_clause, [], [f4294, f1095, f636, f420])).
fof(f4294, plain, (~ (e0 = op(e3, e1)) | (~ spl30_60 | spl30_134)), inference(backward_demodulation, [], [f1097, f638])).
fof(f4296, plain, (~ spl30_44 | ~ spl30_60), inference(avatar_split_clause, [], [f4291, f636, f568])).
fof(f4291, plain, (~ (e3 = op(e1, e1)) | ~ spl30_60), inference(backward_demodulation, [], [f169, f638])).
fof(f4287, plain, (~ spl30_59 | ~ spl30_63), inference(avatar_split_clause, [], [f4282, f649, f632])).
fof(f4282, plain, (~ (e2 = op(e0, e1)) | ~ spl30_63), inference(backward_demodulation, [], [f187, f651])).
fof(f4278, plain, (spl30_60 | ~ spl30_50 | ~ spl30_89), inference(avatar_split_clause, [], [f4277, f852, f594, f636])).
fof(f4277, plain, ((e3 = op(e0, e1)) | (~ spl30_50 | ~ spl30_89)), inference(forward_demodulation, [], [f854, f596])).
fof(f854, plain, ((e3 = op(e0, op(e0, e3))) | ~ spl30_89), inference(avatar_component_clause, [], [f852])).
fof(f4276, plain, (spl30_63 | ~ spl30_53 | ~ spl30_90), inference(avatar_split_clause, [], [f4275, f857, f607, f649])).
fof(f4275, plain, ((op(e0, e0) = e2) | (~ spl30_53 | ~ spl30_90)), inference(forward_demodulation, [], [f859, f609])).
fof(f4274, plain, (spl30_48 | ~ spl30_14 | ~ spl30_103), inference(avatar_split_clause, [], [f4273, f938, f441, f585])).
fof(f4273, plain, ((e3 = op(e1, e0)) | (~ spl30_14 | ~ spl30_103)), inference(forward_demodulation, [], [f939, f443])).
fof(f939, plain, ((e3 = op(op(e3, e0), e0)) | ~ spl30_103), inference(avatar_component_clause, [], [f938])).
fof(f4271, plain, (~ spl30_33 | ~ spl30_50 | spl30_128), inference(avatar_split_clause, [], [f4270, f1065, f594, f522])).
fof(f4270, plain, (~ (e0 = op(e1, e3)) | (~ spl30_50 | spl30_128)), inference(forward_demodulation, [], [f1067, f596])).
fof(f4249, plain, (~ spl30_60 | ~ spl30_156 | spl30_160), inference(avatar_split_clause, [], [f4248, f1232, f1210, f636])).
fof(f4248, plain, (~ (e3 = op(e0, e1)) | (~ spl30_156 | spl30_160)), inference(forward_demodulation, [], [f1234, f1211])).
fof(f1211, plain, ((e0 = op(op(e1, e1), e1)) | ~ spl30_156), inference(avatar_component_clause, [], [f1210])).
fof(f4245, plain, (spl30_33 | ~ spl30_14 | ~ spl30_104), inference(avatar_split_clause, [], [f4019, f943, f441, f522])).
fof(f4019, plain, ((e0 = op(e1, e3)) | (~ spl30_14 | ~ spl30_104)), inference(forward_demodulation, [], [f945, f443])).
fof(f4244, plain, (~ spl30_34 | ~ spl30_38), inference(avatar_split_clause, [], [f4243, f543, f526])).
fof(f4243, plain, (~ (e1 = op(e1, e3)) | ~ spl30_38), inference(forward_demodulation, [], [f198, f545])).
fof(f545, plain, ((e1 = op(e1, e2)) | ~ spl30_38), inference(avatar_component_clause, [], [f543])).
fof(f4241, plain, (~ spl30_9 | ~ spl30_14 | spl30_68), inference(avatar_contradiction_clause, [], [f4240])).
fof(f4240, plain, ($false | (~ spl30_9 | ~ spl30_14 | spl30_68)), inference(subsumption_resolution, [], [f4239, f443])).
fof(f4239, plain, (~ (e1 = op(e3, e0)) | (~ spl30_9 | spl30_68)), inference(forward_demodulation, [], [f705, f422])).
fof(f4227, plain, (~ spl30_9 | ~ spl30_44 | ~ spl30_59 | spl30_155), inference(avatar_contradiction_clause, [], [f4226])).
fof(f4226, plain, ($false | (~ spl30_9 | ~ spl30_44 | ~ spl30_59 | spl30_155)), inference(subsumption_resolution, [], [f4225, f634])).
fof(f4225, plain, (~ (e2 = op(e0, e1)) | (~ spl30_9 | ~ spl30_44 | spl30_155)), inference(forward_demodulation, [], [f4224, f422])).
fof(f4224, plain, (~ (e2 = op(op(e3, e1), e1)) | (~ spl30_44 | spl30_155)), inference(forward_demodulation, [], [f1208, f570])).
fof(f1208, plain, (~ (e2 = op(op(op(e1, e1), e1), e1)) | spl30_155), inference(avatar_component_clause, [], [f1206])).
fof(f1206, plain, (spl30_155 <=> (e2 = op(op(op(e1, e1), e1), e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_155])])).
fof(f4222, plain, (~ spl30_14 | ~ spl30_47 | ~ spl30_64 | spl30_158), inference(avatar_contradiction_clause, [], [f4221])).
fof(f4221, plain, ($false | (~ spl30_14 | ~ spl30_47 | ~ spl30_64 | spl30_158)), inference(subsumption_resolution, [], [f4220, f583])).
fof(f4220, plain, (~ (e2 = op(e1, e0)) | (~ spl30_14 | ~ spl30_64 | spl30_158)), inference(forward_demodulation, [], [f4219, f443])).
fof(f4219, plain, (~ (e2 = op(op(e3, e0), e0)) | (~ spl30_64 | spl30_158)), inference(forward_demodulation, [], [f1222, f655])).
fof(f1222, plain, (~ (e2 = op(op(op(e0, e0), e0), e0)) | spl30_158), inference(avatar_component_clause, [], [f1220])).
fof(f1220, plain, (spl30_158 <=> (e2 = op(op(op(e0, e0), e0), e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_158])])).
fof(f4193, plain, (~ spl30_5 | ~ spl30_24 | spl30_149), inference(avatar_split_clause, [], [f4187, f1176, f483, f403])).
fof(f4187, plain, (~ (e0 = op(e3, e2)) | (~ spl30_24 | spl30_149)), inference(backward_demodulation, [], [f1178, f485])).
fof(f4189, plain, (~ spl30_19 | ~ spl30_24 | spl30_74), inference(avatar_contradiction_clause, [], [f4188])).
fof(f4188, plain, ($false | (~ spl30_19 | ~ spl30_24 | spl30_74)), inference(subsumption_resolution, [], [f4183, f464])).
fof(f4183, plain, (~ (e2 = op(e2, e3)) | (~ spl30_24 | spl30_74)), inference(backward_demodulation, [], [f746, f485])).
fof(f4174, plain, (~ spl30_38 | spl30_43 | ~ spl30_82), inference(avatar_contradiction_clause, [], [f4173])).
fof(f4173, plain, ($false | (~ spl30_38 | spl30_43 | ~ spl30_82)), inference(subsumption_resolution, [], [f4170, f565])).
fof(f565, plain, (~ (e2 = op(e1, e1)) | spl30_43), inference(avatar_component_clause, [], [f564])).
fof(f4170, plain, ((e2 = op(e1, e1)) | (~ spl30_38 | ~ spl30_82)), inference(backward_demodulation, [], [f803, f545])).
fof(f4172, plain, (~ spl30_22 | ~ spl30_38), inference(avatar_split_clause, [], [f4166, f543, f475])).
fof(f4166, plain, (~ (e1 = op(e2, e2)) | ~ spl30_38), inference(backward_demodulation, [], [f178, f545])).
fof(f4165, plain, (~ spl30_9 | ~ spl30_44 | spl30_156), inference(avatar_split_clause, [], [f4160, f1210, f568, f420])).
fof(f4160, plain, (~ (e0 = op(e3, e1)) | (~ spl30_44 | spl30_156)), inference(backward_demodulation, [], [f1212, f570])).
fof(f4164, plain, (~ spl30_11 | ~ spl30_44 | spl30_144), inference(avatar_split_clause, [], [f4159, f1152, f568, f428])).
fof(f4159, plain, (~ (e2 = op(e3, e1)) | (~ spl30_44 | spl30_144)), inference(backward_demodulation, [], [f1154, f570])).
fof(f4162, plain, (~ spl30_28 | ~ spl30_44), inference(avatar_split_clause, [], [f4154, f568, f500])).
fof(f4154, plain, (~ (e3 = op(e2, e1)) | ~ spl30_44), inference(backward_demodulation, [], [f172, f570])).
fof(f4145, plain, (~ spl30_5 | ~ spl30_53), inference(avatar_split_clause, [], [f4139, f607, f403])).
fof(f4139, plain, (~ (e0 = op(e3, e2)) | ~ spl30_53), inference(backward_demodulation, [], [f177, f609])).
fof(f4124, plain, (~ spl30_14 | ~ spl30_64 | spl30_159), inference(avatar_contradiction_clause, [], [f4123])).
fof(f4123, plain, ($false | (~ spl30_14 | ~ spl30_64 | spl30_159)), inference(subsumption_resolution, [], [f4115, f443])).
fof(f4115, plain, (~ (e1 = op(e3, e0)) | (~ spl30_64 | spl30_159)), inference(backward_demodulation, [], [f1226, f655])).
fof(f4104, plain, (spl30_53 | ~ spl30_29 | ~ spl30_115), inference(avatar_split_clause, [], [f4103, f999, f505, f607])).
fof(f4103, plain, ((e0 = op(e0, e2)) | (~ spl30_29 | ~ spl30_115)), inference(forward_demodulation, [], [f1001, f507])).
fof(f4087, plain, (~ spl30_24 | ~ spl30_19 | spl30_73), inference(avatar_split_clause, [], [f4060, f740, f462, f483])).
fof(f4060, plain, (~ (e3 = op(e2, e2)) | (~ spl30_19 | spl30_73)), inference(backward_demodulation, [], [f741, f464])).
fof(f4039, plain, (~ spl30_27 | ~ spl30_22 | spl30_74), inference(avatar_split_clause, [], [f4038, f745, f475, f496])).
fof(f4038, plain, (~ (e2 = op(e2, e1)) | (~ spl30_22 | spl30_74)), inference(forward_demodulation, [], [f746, f477])).
fof(f3976, plain, (~ spl30_42 | spl30_83), inference(avatar_contradiction_clause, [], [f3975])).
fof(f3975, plain, ($false | (~ spl30_42 | spl30_83)), inference(subsumption_resolution, [], [f3972, f562])).
fof(f3972, plain, (~ (e1 = op(e1, e1)) | (~ spl30_42 | spl30_83)), inference(backward_demodulation, [], [f807, f562])).
fof(f3957, plain, (~ spl30_2 | ~ spl30_50), inference(avatar_split_clause, [], [f3952, f594, f390])).
fof(f3952, plain, (~ (e1 = op(e3, e3)) | ~ spl30_50), inference(backward_demodulation, [], [f183, f596])).
fof(f3951, plain, (spl30_5 | ~ spl30_56 | ~ spl30_131), inference(avatar_split_clause, [], [f3948, f1080, f619, f403])).
fof(f3948, plain, ((e0 = op(e3, e2)) | (~ spl30_56 | ~ spl30_131)), inference(backward_demodulation, [], [f1081, f621])).
fof(f3950, plain, (~ spl30_51 | ~ spl30_56 | spl30_90), inference(avatar_split_clause, [], [f3947, f857, f619, f598])).
fof(f3947, plain, (~ (e2 = op(e0, e3)) | (~ spl30_56 | spl30_90)), inference(backward_demodulation, [], [f858, f621])).
fof(f3944, plain, (~ spl30_53 | ~ spl30_57), inference(avatar_split_clause, [], [f3936, f624, f607])).
fof(f3936, plain, (~ (e0 = op(e0, e2)) | ~ spl30_57), inference(backward_demodulation, [], [f190, f626])).
fof(f3943, plain, (~ spl30_9 | ~ spl30_57), inference(avatar_split_clause, [], [f3935, f624, f420])).
fof(f3935, plain, (~ (e0 = op(e3, e1)) | ~ spl30_57), inference(backward_demodulation, [], [f171, f626])).
fof(f3929, plain, (~ spl30_53 | ~ spl30_63 | spl30_92), inference(avatar_split_clause, [], [f3922, f867, f649, f607])).
fof(f3922, plain, (~ (e0 = op(e0, e2)) | (~ spl30_63 | spl30_92)), inference(backward_demodulation, [], [f868, f651])).
fof(f3927, plain, (~ spl30_15 | ~ spl30_63), inference(avatar_split_clause, [], [f3917, f649, f445])).
fof(f3917, plain, (~ (e2 = op(e3, e0)) | ~ spl30_63), inference(backward_demodulation, [], [f165, f651])).
fof(f3907, plain, (~ spl30_60 | ~ spl30_33 | spl30_118), inference(avatar_split_clause, [], [f3840, f1014, f522, f636])).
fof(f3840, plain, (~ (e3 = op(e0, e1)) | (~ spl30_33 | spl30_118)), inference(backward_demodulation, [], [f1015, f524])).
fof(f1015, plain, (~ (e3 = op(op(e1, e3), e1)) | spl30_118), inference(avatar_component_clause, [], [f1014])).
fof(f3861, plain, (~ spl30_6 | ~ spl30_22), inference(avatar_split_clause, [], [f3854, f475, f407])).
fof(f3854, plain, (~ (e1 = op(e3, e2)) | ~ spl30_22), inference(backward_demodulation, [], [f180, f477])).
fof(f3850, plain, (~ spl30_13 | ~ spl30_29), inference(avatar_split_clause, [], [f3843, f505, f437])).
fof(f3843, plain, (~ (e0 = op(e3, e0)) | ~ spl30_29), inference(backward_demodulation, [], [f168, f507])).
fof(f3800, plain, (~ spl30_3 | ~ spl30_51), inference(avatar_split_clause, [], [f3795, f598, f394])).
fof(f3795, plain, (~ (e2 = op(e3, e3)) | ~ spl30_51), inference(backward_demodulation, [], [f183, f600])).
fof(f3791, plain, (~ spl30_49 | ~ spl30_53), inference(avatar_split_clause, [], [f3784, f607, f590])).
fof(f3784, plain, (~ (e0 = op(e0, e3)) | ~ spl30_53), inference(backward_demodulation, [], [f192, f609])).
fof(f3790, plain, (~ spl30_21 | ~ spl30_53), inference(avatar_split_clause, [], [f3782, f607, f471])).
fof(f3782, plain, (~ (e0 = op(e2, e2)) | ~ spl30_53), inference(backward_demodulation, [], [f176, f609])).
fof(f3780, plain, (spl30_22 | ~ spl30_27 | ~ spl30_75), inference(avatar_split_clause, [], [f3779, f750, f496, f475])).
fof(f3779, plain, ((e1 = op(e2, e2)) | (~ spl30_27 | ~ spl30_75)), inference(forward_demodulation, [], [f752, f498])).
fof(f3778, plain, (~ spl30_58 | spl30_91), inference(avatar_contradiction_clause, [], [f3777])).
fof(f3777, plain, ($false | (~ spl30_58 | spl30_91)), inference(subsumption_resolution, [], [f3776, f630])).
fof(f3776, plain, (~ (e1 = op(e0, e1)) | (~ spl30_58 | spl30_91)), inference(forward_demodulation, [], [f863, f630])).
fof(f3771, plain, (spl30_46 | ~ spl30_58 | ~ spl30_135), inference(avatar_split_clause, [], [f3770, f1100, f628, f577])).
fof(f3770, plain, ((e1 = op(e1, e0)) | (~ spl30_58 | ~ spl30_135)), inference(forward_demodulation, [], [f1102, f630])).
fof(f3767, plain, (~ spl30_51 | ~ spl30_151 | spl30_157), inference(avatar_split_clause, [], [f3765, f1215, f1185, f598])).
fof(f3765, plain, (~ (e2 = op(e0, e3)) | (~ spl30_151 | spl30_157)), inference(backward_demodulation, [], [f1217, f1186])).
fof(f1186, plain, ((e0 = op(op(e3, e3), e3)) | ~ spl30_151), inference(avatar_component_clause, [], [f1185])).
fof(f3758, plain, (~ spl30_36 | ~ spl30_20), inference(avatar_split_clause, [], [f3730, f466, f534])).
fof(f3730, plain, (~ (e3 = op(e1, e3)) | ~ spl30_20), inference(forward_demodulation, [], [f184, f468])).
fof(f3743, plain, (~ spl30_12 | spl30_100), inference(avatar_contradiction_clause, [], [f3742])).
fof(f3742, plain, ($false | (~ spl30_12 | spl30_100)), inference(subsumption_resolution, [], [f3741, f434])).
fof(f3741, plain, (~ (e3 = op(e3, e1)) | (~ spl30_12 | spl30_100)), inference(forward_demodulation, [], [f925, f434])).
fof(f925, plain, (~ (e3 = op(op(e3, e1), e1)) | spl30_100), inference(avatar_component_clause, [], [f923])).
fof(f3722, plain, (spl30_44 | ~ spl30_34 | ~ spl30_118), inference(avatar_split_clause, [], [f3721, f1014, f526, f568])).
fof(f3721, plain, ((e3 = op(e1, e1)) | (~ spl30_34 | ~ spl30_118)), inference(forward_demodulation, [], [f1016, f528])).
fof(f3703, plain, (~ spl30_9 | ~ spl30_12), inference(avatar_contradiction_clause, [], [f3702])).
fof(f3702, plain, ($false | (~ spl30_9 | ~ spl30_12)), inference(subsumption_resolution, [], [f3701, f213])).
fof(f3701, plain, ((e0 = e3) | (~ spl30_9 | ~ spl30_12)), inference(forward_demodulation, [], [f434, f422])).
fof(f3698, plain, (~ spl30_13 | spl30_69), inference(avatar_contradiction_clause, [], [f3697])).
fof(f3697, plain, ($false | (~ spl30_13 | spl30_69)), inference(subsumption_resolution, [], [f3692, f439])).
fof(f3692, plain, (~ (e0 = op(e3, e0)) | (~ spl30_13 | spl30_69)), inference(backward_demodulation, [], [f710, f439])).
fof(f3674, plain, (~ spl30_33 | ~ spl30_34), inference(avatar_contradiction_clause, [], [f3673])).
fof(f3673, plain, ($false | (~ spl30_33 | ~ spl30_34)), inference(subsumption_resolution, [], [f3672, f211])).
fof(f211, plain, ~ (e0 = e1), inference(cnf_transformation, [], [f4])).
fof(f3672, plain, ((e0 = e1) | (~ spl30_33 | ~ spl30_34)), inference(forward_demodulation, [], [f528, f524])).
fof(f3616, plain, (~ spl30_50 | ~ spl30_9 | spl30_101), inference(avatar_split_clause, [], [f3615, f928, f420, f594])).
fof(f3615, plain, (~ (e1 = op(e0, e3)) | (~ spl30_9 | spl30_101)), inference(forward_demodulation, [], [f929, f422])).
fof(f3611, plain, (spl30_60 | ~ spl30_9 | ~ spl30_100), inference(avatar_split_clause, [], [f3274, f923, f420, f636])).
fof(f3274, plain, ((e3 = op(e0, e1)) | (~ spl30_9 | ~ spl30_100)), inference(backward_demodulation, [], [f924, f422])).
fof(f3609, plain, (~ spl30_49 | ~ spl30_33), inference(avatar_split_clause, [], [f3608, f522, f590])).
fof(f3608, plain, (~ (e0 = op(e0, e3)) | ~ spl30_33), inference(forward_demodulation, [], [f181, f524])).
fof(f3600, plain, (~ spl30_13 | ~ spl30_9), inference(avatar_split_clause, [], [f3425, f420, f437])).
fof(f3425, plain, (~ (e0 = op(e3, e0)) | ~ spl30_9), inference(forward_demodulation, [], [f205, f422])).
fof(f3593, plain, (~ spl30_31 | spl30_114), inference(avatar_contradiction_clause, [], [f3592])).
fof(f3592, plain, ($false | (~ spl30_31 | spl30_114)), inference(subsumption_resolution, [], [f3587, f515])).
fof(f3587, plain, (~ (e2 = op(e2, e0)) | (~ spl30_31 | spl30_114)), inference(backward_demodulation, [], [f996, f515])).
fof(f3589, plain, (~ spl30_27 | ~ spl30_31), inference(avatar_split_clause, [], [f3584, f513, f496])).
fof(f3584, plain, (~ (e2 = op(e2, e1)) | ~ spl30_31), inference(backward_demodulation, [], [f199, f515])).
fof(f3583, plain, (~ spl30_27 | ~ spl30_43), inference(avatar_split_clause, [], [f3578, f564, f496])).
fof(f3578, plain, (~ (e2 = op(e2, e1)) | ~ spl30_43), inference(backward_demodulation, [], [f172, f566])).
fof(f3555, plain, (~ spl30_31 | ~ spl30_55 | spl30_132), inference(avatar_split_clause, [], [f3362, f1085, f615, f513])).
fof(f3362, plain, (~ (e2 = op(e2, e0)) | (~ spl30_55 | spl30_132)), inference(backward_demodulation, [], [f1086, f617])).
fof(f3549, plain, (~ spl30_31 | ~ spl30_21 | spl30_74), inference(avatar_split_clause, [], [f3484, f745, f471, f513])).
fof(f3484, plain, (~ (e2 = op(e2, e0)) | (~ spl30_21 | spl30_74)), inference(backward_demodulation, [], [f746, f473])).
fof(f3541, plain, (~ spl30_6 | ~ spl30_40 | spl30_120), inference(avatar_contradiction_clause, [], [f3540])).
fof(f3540, plain, ($false | (~ spl30_6 | ~ spl30_40 | spl30_120)), inference(subsumption_resolution, [], [f3539, f409])).
fof(f3539, plain, (~ (e1 = op(e3, e2)) | (~ spl30_40 | spl30_120)), inference(forward_demodulation, [], [f1026, f553])).
fof(f1026, plain, (~ (e1 = op(op(e1, e2), e2)) | spl30_120), inference(avatar_component_clause, [], [f1024])).
fof(f3538, plain, (~ spl30_9 | ~ spl30_40 | ~ spl30_121), inference(avatar_contradiction_clause, [], [f3537])).
fof(f3537, plain, ($false | (~ spl30_9 | ~ spl30_40 | ~ spl30_121)), inference(subsumption_resolution, [], [f3536, f212])).
fof(f212, plain, ~ (e0 = e2), inference(cnf_transformation, [], [f4])).
fof(f3536, plain, ((e0 = e2) | (~ spl30_9 | ~ spl30_40 | ~ spl30_121)), inference(forward_demodulation, [], [f3535, f422])).
fof(f3535, plain, ((e2 = op(e3, e1)) | (~ spl30_40 | ~ spl30_121)), inference(forward_demodulation, [], [f1031, f553])).
fof(f3496, plain, (~ spl30_20 | spl30_73), inference(avatar_contradiction_clause, [], [f3495])).
fof(f3495, plain, ($false | (~ spl30_20 | spl30_73)), inference(subsumption_resolution, [], [f3493, f468])).
fof(f3493, plain, (~ (e3 = op(e2, e3)) | (~ spl30_20 | spl30_73)), inference(backward_demodulation, [], [f741, f468])).
fof(f3488, plain, (~ spl30_17 | ~ spl30_21), inference(avatar_split_clause, [], [f3480, f471, f454])).
fof(f3480, plain, (~ (e0 = op(e2, e3)) | ~ spl30_21), inference(backward_demodulation, [], [f204, f473])).
fof(f3462, plain, (~ spl30_30 | ~ spl30_47 | spl30_125), inference(avatar_contradiction_clause, [], [f3461])).
fof(f3461, plain, ($false | (~ spl30_30 | ~ spl30_47 | spl30_125)), inference(subsumption_resolution, [], [f3456, f511])).
fof(f3456, plain, (~ (e1 = op(e2, e0)) | (~ spl30_47 | spl30_125)), inference(backward_demodulation, [], [f1052, f583])).
fof(f3433, plain, (spl30_47 | ~ spl30_30 | ~ spl30_114), inference(avatar_split_clause, [], [f3405, f994, f509, f581])).
fof(f3405, plain, ((e2 = op(e1, e0)) | (~ spl30_30 | ~ spl30_114)), inference(backward_demodulation, [], [f995, f511])).
fof(f995, plain, ((e2 = op(op(e2, e0), e0)) | ~ spl30_114), inference(avatar_component_clause, [], [f994])).
fof(f3430, plain, (~ spl30_18 | ~ spl30_50), inference(avatar_split_clause, [], [f3429, f594, f458])).
fof(f3429, plain, (~ (e1 = op(e2, e3)) | ~ spl30_50), inference(forward_demodulation, [], [f182, f596])).
fof(f3422, plain, (~ spl30_6 | ~ spl30_24 | spl30_140), inference(avatar_contradiction_clause, [], [f3421])).
fof(f3421, plain, ($false | (~ spl30_6 | ~ spl30_24 | spl30_140)), inference(subsumption_resolution, [], [f3414, f409])).
fof(f3414, plain, (~ (e1 = op(e3, e2)) | (~ spl30_24 | spl30_140)), inference(backward_demodulation, [], [f1136, f485])).
fof(f3420, plain, (~ spl30_6 | ~ spl30_24 | ~ spl30_37 | spl30_139), inference(avatar_contradiction_clause, [], [f3419])).
fof(f3419, plain, ($false | (~ spl30_6 | ~ spl30_24 | ~ spl30_37 | spl30_139)), inference(subsumption_resolution, [], [f3418, f541])).
fof(f3418, plain, (~ (e0 = op(e1, e2)) | (~ spl30_6 | ~ spl30_24 | spl30_139)), inference(forward_demodulation, [], [f3413, f409])).
fof(f3413, plain, (~ (e0 = op(op(e3, e2), e2)) | (~ spl30_24 | spl30_139)), inference(backward_demodulation, [], [f1132, f485])).
fof(f1132, plain, (~ (e0 = op(op(op(e2, e2), e2), e2)) | spl30_139), inference(avatar_component_clause, [], [f1130])).
fof(f1130, plain, (spl30_139 <=> (e0 = op(op(op(e2, e2), e2), e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_139])])).
fof(f3411, plain, (~ spl30_30 | ~ spl30_37 | spl30_115), inference(avatar_contradiction_clause, [], [f3410])).
fof(f3410, plain, ($false | (~ spl30_30 | ~ spl30_37 | spl30_115)), inference(subsumption_resolution, [], [f3406, f541])).
fof(f3406, plain, (~ (e0 = op(e1, e2)) | (~ spl30_30 | spl30_115)), inference(backward_demodulation, [], [f1000, f511])).
fof(f3407, plain, (~ spl30_22 | ~ spl30_30), inference(avatar_split_clause, [], [f3401, f509, f475])).
fof(f3401, plain, (~ (e1 = op(e2, e2)) | ~ spl30_30), inference(backward_demodulation, [], [f200, f511])).
fof(f3366, plain, (spl30_21 | ~ spl30_55 | ~ spl30_131), inference(avatar_contradiction_clause, [], [f3365])).
fof(f3365, plain, ($false | (spl30_21 | ~ spl30_55 | ~ spl30_131)), inference(subsumption_resolution, [], [f3361, f472])).
fof(f3361, plain, ((e0 = op(e2, e2)) | (~ spl30_55 | ~ spl30_131)), inference(backward_demodulation, [], [f1081, f617])).
fof(f3333, plain, (~ spl30_32 | ~ spl30_17 | spl30_73), inference(avatar_split_clause, [], [f3332, f740, f454, f517])).
fof(f3332, plain, (~ (e3 = op(e2, e0)) | (~ spl30_17 | spl30_73)), inference(forward_demodulation, [], [f741, f456])).
fof(f3314, plain, (spl30_50 | ~ spl30_60 | ~ spl30_91), inference(avatar_split_clause, [], [f3234, f862, f636, f594])).
fof(f3234, plain, ((e1 = op(e0, e3)) | (~ spl30_60 | ~ spl30_91)), inference(backward_demodulation, [], [f864, f638])).
fof(f3288, plain, (~ spl30_6 | ~ spl30_9 | ~ spl30_67), inference(avatar_contradiction_clause, [], [f3287])).
fof(f3287, plain, ($false | (~ spl30_6 | ~ spl30_9 | ~ spl30_67)), inference(subsumption_resolution, [], [f3286, f212])).
fof(f3286, plain, ((e0 = e2) | (~ spl30_6 | ~ spl30_9 | ~ spl30_67)), inference(forward_demodulation, [], [f3282, f422])).
fof(f3282, plain, ((e2 = op(e3, e1)) | (~ spl30_6 | ~ spl30_67)), inference(backward_demodulation, [], [f701, f409])).
fof(f3285, plain, (~ spl30_2 | ~ spl30_6), inference(avatar_split_clause, [], [f3279, f407, f390])).
fof(f3279, plain, (~ (e1 = op(e3, e3)) | ~ spl30_6), inference(backward_demodulation, [], [f210, f409])).
fof(f3278, plain, (~ spl30_9 | spl30_14 | ~ spl30_68), inference(avatar_contradiction_clause, [], [f3277])).
fof(f3277, plain, ($false | (~ spl30_9 | spl30_14 | ~ spl30_68)), inference(subsumption_resolution, [], [f3273, f442])).
fof(f442, plain, (~ (e1 = op(e3, e0)) | spl30_14), inference(avatar_component_clause, [], [f441])).
fof(f3273, plain, ((e1 = op(e3, e0)) | (~ spl30_9 | ~ spl30_68)), inference(backward_demodulation, [], [f706, f422])).
fof(f3264, plain, (~ spl30_34 | spl30_117), inference(avatar_contradiction_clause, [], [f3263])).
fof(f3263, plain, ($false | (~ spl30_34 | spl30_117)), inference(subsumption_resolution, [], [f3259, f528])).
fof(f3259, plain, (~ (e1 = op(e1, e3)) | (~ spl30_34 | spl30_117)), inference(backward_demodulation, [], [f1011, f528])).
fof(f3261, plain, (~ spl30_2 | ~ spl30_34), inference(avatar_split_clause, [], [f3255, f526, f390])).
fof(f3255, plain, (~ (e1 = op(e3, e3)) | ~ spl30_34), inference(backward_demodulation, [], [f185, f528])).
fof(f3254, plain, (~ spl30_39 | ~ spl30_40), inference(avatar_contradiction_clause, [], [f3253])).
fof(f3253, plain, ($false | (~ spl30_39 | ~ spl30_40)), inference(subsumption_resolution, [], [f3252, f216])).
fof(f3252, plain, ((e2 = e3) | (~ spl30_39 | ~ spl30_40)), inference(forward_demodulation, [], [f553, f549])).
fof(f3238, plain, (~ spl30_56 | ~ spl30_60), inference(avatar_split_clause, [], [f3233, f636, f619])).
fof(f3233, plain, (~ (e3 = op(e0, e2)) | ~ spl30_60), inference(backward_demodulation, [], [f190, f638])).
fof(f3231, plain, (~ spl30_5 | ~ spl30_15 | spl30_69), inference(avatar_split_clause, [], [f3230, f709, f445, f403])).
fof(f3230, plain, (~ (e0 = op(e3, e2)) | (~ spl30_15 | spl30_69)), inference(forward_demodulation, [], [f710, f447])).
fof(f3228, plain, (~ spl30_56 | ~ spl30_51 | spl30_89), inference(avatar_split_clause, [], [f3227, f852, f598, f619])).
fof(f3227, plain, (~ (e3 = op(e0, e2)) | (~ spl30_51 | spl30_89)), inference(forward_demodulation, [], [f853, f600])).
fof(f3202, plain, (~ spl30_27 | spl30_111), inference(avatar_contradiction_clause, [], [f3201])).
fof(f3201, plain, ($false | (~ spl30_27 | spl30_111)), inference(subsumption_resolution, [], [f3200, f498])).
fof(f3200, plain, (~ (e2 = op(e2, e1)) | (~ spl30_27 | spl30_111)), inference(forward_demodulation, [], [f981, f498])).
fof(f3181, plain, (~ spl30_5 | ~ spl30_56 | spl30_131), inference(avatar_split_clause, [], [f3176, f1080, f619, f403])).
fof(f3176, plain, (~ (e0 = op(e3, e2)) | (~ spl30_56 | spl30_131)), inference(backward_demodulation, [], [f1082, f621])).
fof(f3158, plain, (spl30_56 | ~ spl30_51 | ~ spl30_89), inference(avatar_split_clause, [], [f3157, f852, f598, f619])).
fof(f3157, plain, ((e3 = op(e0, e2)) | (~ spl30_51 | ~ spl30_89)), inference(forward_demodulation, [], [f854, f600])).
fof(f3156, plain, (spl30_57 | ~ spl30_62 | ~ spl30_92), inference(avatar_split_clause, [], [f3155, f867, f645, f624])).
fof(f3155, plain, ((e0 = op(e0, e1)) | (~ spl30_62 | ~ spl30_92)), inference(forward_demodulation, [], [f869, f647])).
fof(f647, plain, ((op(e0, e0) = e1) | ~ spl30_62), inference(avatar_component_clause, [], [f645])).
fof(f3133, plain, (~ spl30_11 | ~ spl30_27), inference(avatar_split_clause, [], [f3132, f496, f428])).
fof(f3132, plain, (~ (e2 = op(e3, e1)) | ~ spl30_27), inference(forward_demodulation, [], [f174, f498])).
fof(f3108, plain, (~ spl30_19 | ~ spl30_27), inference(avatar_split_clause, [], [f3104, f496, f462])).
fof(f3104, plain, (~ (e2 = op(e2, e3)) | ~ spl30_27), inference(backward_demodulation, [], [f203, f498])).
fof(f3102, plain, (spl30_27 | ~ spl30_39 | ~ spl30_121), inference(avatar_split_clause, [], [f3098, f1029, f547, f496])).
fof(f3098, plain, ((e2 = op(e2, e1)) | (~ spl30_39 | ~ spl30_121)), inference(backward_demodulation, [], [f1031, f549])).
fof(f3101, plain, (~ spl30_22 | ~ spl30_39 | spl30_120), inference(avatar_split_clause, [], [f3097, f1024, f547, f475])).
fof(f3097, plain, (~ (e1 = op(e2, e2)) | (~ spl30_39 | spl30_120)), inference(backward_demodulation, [], [f1026, f549])).
fof(f3100, plain, (~ spl30_39 | spl30_82), inference(avatar_contradiction_clause, [], [f3099])).
fof(f3099, plain, ($false | (~ spl30_39 | spl30_82)), inference(subsumption_resolution, [], [f3096, f549])).
fof(f3096, plain, (~ (e2 = op(e1, e2)) | (~ spl30_39 | spl30_82)), inference(backward_demodulation, [], [f802, f549])).
fof(f3070, plain, (~ spl30_53 | ~ spl30_62 | ~ spl30_90), inference(avatar_contradiction_clause, [], [f3069])).
fof(f3069, plain, ($false | (~ spl30_53 | ~ spl30_62 | ~ spl30_90)), inference(subsumption_resolution, [], [f3068, f214])).
fof(f3068, plain, ((e1 = e2) | (~ spl30_53 | ~ spl30_62 | ~ spl30_90)), inference(forward_demodulation, [], [f3063, f647])).
fof(f3063, plain, ((op(e0, e0) = e2) | (~ spl30_53 | ~ spl30_90)), inference(backward_demodulation, [], [f859, f609])).
fof(f3054, plain, (spl30_17 | ~ spl30_15 | ~ spl30_104), inference(avatar_split_clause, [], [f3053, f943, f445, f454])).
fof(f3053, plain, ((e0 = op(e2, e3)) | (~ spl30_15 | ~ spl30_104)), inference(forward_demodulation, [], [f945, f447])).
fof(f3038, plain, (~ spl30_19 | spl30_106), inference(avatar_contradiction_clause, [], [f3037])).
fof(f3037, plain, ($false | (~ spl30_19 | spl30_106)), inference(subsumption_resolution, [], [f3036, f464])).
fof(f3036, plain, (~ (e2 = op(e2, e3)) | (~ spl30_19 | spl30_106)), inference(forward_demodulation, [], [f955, f464])).
fof(f3035, plain, (~ spl30_19 | ~ spl30_21 | ~ spl30_107), inference(avatar_contradiction_clause, [], [f3034])).
fof(f3034, plain, ($false | (~ spl30_19 | ~ spl30_21 | ~ spl30_107)), inference(subsumption_resolution, [], [f3033, f213])).
fof(f3033, plain, ((e0 = e3) | (~ spl30_19 | ~ spl30_21 | ~ spl30_107)), inference(forward_demodulation, [], [f3032, f473])).
fof(f3032, plain, ((e3 = op(e2, e2)) | (~ spl30_19 | ~ spl30_107)), inference(forward_demodulation, [], [f960, f464])).
fof(f3024, plain, (~ spl30_19 | ~ spl30_21 | ~ spl30_73), inference(avatar_contradiction_clause, [], [f3023])).
fof(f3023, plain, ($false | (~ spl30_19 | ~ spl30_21 | ~ spl30_73)), inference(subsumption_resolution, [], [f3022, f213])).
fof(f3022, plain, ((e0 = e3) | (~ spl30_19 | ~ spl30_21 | ~ spl30_73)), inference(forward_demodulation, [], [f3019, f473])).
fof(f3019, plain, ((e3 = op(e2, e2)) | (~ spl30_19 | ~ spl30_73)), inference(backward_demodulation, [], [f742, f464])).
fof(f3012, plain, (~ spl30_21 | spl30_31 | ~ spl30_74), inference(avatar_contradiction_clause, [], [f3011])).
fof(f3011, plain, ($false | (~ spl30_21 | spl30_31 | ~ spl30_74)), inference(subsumption_resolution, [], [f3005, f514])).
fof(f3005, plain, ((e2 = op(e2, e0)) | (~ spl30_21 | ~ spl30_74)), inference(backward_demodulation, [], [f747, f473])).
fof(f3000, plain, (~ spl30_49 | ~ spl30_62 | ~ spl30_129), inference(avatar_contradiction_clause, [], [f2999])).
fof(f2999, plain, ($false | (~ spl30_49 | ~ spl30_62 | ~ spl30_129)), inference(subsumption_resolution, [], [f2998, f215])).
fof(f2998, plain, ((e1 = e3) | (~ spl30_49 | ~ spl30_62 | ~ spl30_129)), inference(forward_demodulation, [], [f2995, f647])).
fof(f2995, plain, ((op(e0, e0) = e3) | (~ spl30_49 | ~ spl30_129)), inference(backward_demodulation, [], [f1072, f592])).
fof(f2996, plain, (~ spl30_17 | ~ spl30_49), inference(avatar_split_clause, [], [f2989, f590, f454])).
fof(f2989, plain, (~ (e0 = op(e2, e3)) | ~ spl30_49), inference(backward_demodulation, [], [f182, f592])).
fof(f2986, plain, (~ spl30_55 | spl30_90), inference(avatar_contradiction_clause, [], [f2985])).
fof(f2985, plain, ($false | (~ spl30_55 | spl30_90)), inference(subsumption_resolution, [], [f2982, f617])).
fof(f2982, plain, (~ (e2 = op(e0, e2)) | (~ spl30_55 | spl30_90)), inference(backward_demodulation, [], [f858, f617])).
fof(f2949, plain, (~ spl30_53 | spl30_131), inference(avatar_contradiction_clause, [], [f2948])).
fof(f2948, plain, ($false | (~ spl30_53 | spl30_131)), inference(subsumption_resolution, [], [f2947, f609])).
fof(f2947, plain, (~ (e0 = op(e0, e2)) | (~ spl30_53 | spl30_131)), inference(forward_demodulation, [], [f1082, f609])).
fof(f2926, plain, (~ spl30_9 | spl30_50 | ~ spl30_101), inference(avatar_contradiction_clause, [], [f2925])).
fof(f2925, plain, ($false | (~ spl30_9 | spl30_50 | ~ spl30_101)), inference(subsumption_resolution, [], [f2921, f595])).
fof(f595, plain, (~ (e1 = op(e0, e3)) | spl30_50), inference(avatar_component_clause, [], [f594])).
fof(f2921, plain, ((e1 = op(e0, e3)) | (~ spl30_9 | ~ spl30_101)), inference(backward_demodulation, [], [f930, f422])).
fof(f2924, plain, (~ spl30_9 | ~ spl30_60 | spl30_100), inference(avatar_contradiction_clause, [], [f2923])).
fof(f2923, plain, ($false | (~ spl30_9 | ~ spl30_60 | spl30_100)), inference(subsumption_resolution, [], [f2920, f638])).
fof(f2920, plain, (~ (e3 = op(e0, e1)) | (~ spl30_9 | spl30_100)), inference(backward_demodulation, [], [f925, f422])).
fof(f2874, plain, (~ spl30_10 | ~ spl30_26), inference(avatar_split_clause, [], [f2872, f492, f424])).
fof(f2872, plain, (~ (e1 = op(e3, e1)) | ~ spl30_26), inference(backward_demodulation, [], [f174, f494])).
fof(f2852, plain, (~ spl30_4 | ~ spl30_36), inference(avatar_split_clause, [], [f2851, f534, f398])).
fof(f2851, plain, (~ (e3 = op(e3, e3)) | ~ spl30_36), inference(forward_demodulation, [], [f185, f536])).
fof(f536, plain, ((e3 = op(e1, e3)) | ~ spl30_36), inference(avatar_component_clause, [], [f534])).
fof(f2847, plain, (~ spl30_2 | ~ spl30_36 | spl30_117), inference(avatar_split_clause, [], [f2787, f1009, f534, f390])).
fof(f2787, plain, (~ (e1 = op(e3, e3)) | (~ spl30_36 | spl30_117)), inference(forward_demodulation, [], [f1011, f536])).
fof(f2846, plain, (~ spl30_25 | ~ spl30_26), inference(avatar_contradiction_clause, [], [f2845])).
fof(f2845, plain, ($false | (~ spl30_25 | ~ spl30_26)), inference(subsumption_resolution, [], [f2844, f211])).
fof(f2844, plain, ((e0 = e1) | (~ spl30_25 | ~ spl30_26)), inference(backward_demodulation, [], [f494, f490])).
fof(f2831, plain, (~ spl30_38 | ~ spl30_43 | spl30_121), inference(avatar_contradiction_clause, [], [f2830])).
fof(f2830, plain, ($false | (~ spl30_38 | ~ spl30_43 | spl30_121)), inference(subsumption_resolution, [], [f2829, f566])).
fof(f2829, plain, (~ (e2 = op(e1, e1)) | (~ spl30_38 | spl30_121)), inference(backward_demodulation, [], [f1030, f545])).
fof(f2824, plain, (~ spl30_45 | ~ spl30_62 | spl30_125), inference(avatar_contradiction_clause, [], [f2823])).
fof(f2823, plain, ($false | (~ spl30_45 | ~ spl30_62 | spl30_125)), inference(subsumption_resolution, [], [f2820, f647])).
fof(f2820, plain, (~ (op(e0, e0) = e1) | (~ spl30_45 | spl30_125)), inference(backward_demodulation, [], [f1052, f575])).
fof(f2821, plain, (~ spl30_29 | ~ spl30_45), inference(avatar_split_clause, [], [f2814, f573, f505])).
fof(f2814, plain, (~ (e0 = op(e2, e0)) | ~ spl30_45), inference(backward_demodulation, [], [f166, f575])).
fof(f2798, plain, (~ spl30_62 | ~ spl30_64), inference(avatar_contradiction_clause, [], [f2797])).
fof(f2797, plain, ($false | (~ spl30_62 | ~ spl30_64)), inference(subsumption_resolution, [], [f2796, f215])).
fof(f2796, plain, ((e1 = e3) | (~ spl30_62 | ~ spl30_64)), inference(backward_demodulation, [], [f655, f647])).
fof(f2791, plain, (spl30_38 | ~ spl30_43 | ~ spl30_83), inference(avatar_split_clause, [], [f2790, f806, f564, f543])).
fof(f2790, plain, ((e1 = op(e1, e2)) | (~ spl30_43 | ~ spl30_83)), inference(forward_demodulation, [], [f808, f566])).
fof(f2781, plain, (~ spl30_1 | ~ spl30_18 | ~ spl30_51 | spl30_150), inference(avatar_contradiction_clause, [], [f2780])).
fof(f2780, plain, ($false | (~ spl30_1 | ~ spl30_18 | ~ spl30_51 | spl30_150)), inference(subsumption_resolution, [], [f2779, f460])).
fof(f2779, plain, (~ (e1 = op(e2, e3)) | (~ spl30_1 | ~ spl30_51 | spl30_150)), inference(forward_demodulation, [], [f2658, f600])).
fof(f2658, plain, (~ (e1 = op(op(e0, e3), e3)) | (~ spl30_1 | spl30_150)), inference(forward_demodulation, [], [f1183, f388])).
fof(f1183, plain, (~ (e1 = op(op(op(e3, e3), e3), e3)) | spl30_150), inference(avatar_component_clause, [], [f1181])).
fof(f1181, plain, (spl30_150 <=> (e1 = op(op(op(e3, e3), e3), e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_150])])).
fof(f2762, plain, (~ spl30_32 | ~ spl30_15 | spl30_103), inference(avatar_split_clause, [], [f2495, f938, f445, f517])).
fof(f2495, plain, (~ (e3 = op(e2, e0)) | (~ spl30_15 | spl30_103)), inference(forward_demodulation, [], [f940, f447])).
fof(f2758, plain, (~ spl30_32 | ~ spl30_51 | spl30_129), inference(avatar_split_clause, [], [f2705, f1070, f598, f517])).
fof(f2705, plain, (~ (e3 = op(e2, e0)) | (~ spl30_51 | spl30_129)), inference(backward_demodulation, [], [f1071, f600])).
fof(f1071, plain, (~ (e3 = op(op(e0, e3), e0)) | spl30_129), inference(avatar_component_clause, [], [f1070])).
fof(f2750, plain, (~ spl30_1 | ~ spl30_36 | ~ spl30_117), inference(avatar_contradiction_clause, [], [f2749])).
fof(f2749, plain, ($false | (~ spl30_1 | ~ spl30_36 | ~ spl30_117)), inference(subsumption_resolution, [], [f2748, f211])).
fof(f2748, plain, ((e0 = e1) | (~ spl30_1 | ~ spl30_36 | ~ spl30_117)), inference(forward_demodulation, [], [f2744, f388])).
fof(f2744, plain, ((e1 = op(e3, e3)) | (~ spl30_36 | ~ spl30_117)), inference(backward_demodulation, [], [f1010, f536])).
fof(f2747, plain, (~ spl30_36 | spl30_81), inference(avatar_contradiction_clause, [], [f2746])).
fof(f2746, plain, ($false | (~ spl30_36 | spl30_81)), inference(subsumption_resolution, [], [f2743, f536])).
fof(f2743, plain, (~ (e3 = op(e1, e3)) | (~ spl30_36 | spl30_81)), inference(backward_demodulation, [], [f797, f536])).
fof(f2728, plain, (~ spl30_38 | ~ spl30_43 | spl30_83), inference(avatar_split_clause, [], [f2723, f806, f564, f543])).
fof(f2723, plain, (~ (e1 = op(e1, e2)) | (~ spl30_43 | spl30_83)), inference(backward_demodulation, [], [f807, f566])).
fof(f2718, plain, (~ spl30_38 | ~ spl30_46), inference(avatar_split_clause, [], [f2712, f577, f543])).
fof(f2712, plain, (~ (e1 = op(e1, e2)) | ~ spl30_46), inference(backward_demodulation, [], [f194, f579])).
fof(f2654, plain, (~ spl30_38 | ~ spl30_41 | ~ spl30_121), inference(avatar_contradiction_clause, [], [f2653])).
fof(f2653, plain, ($false | (~ spl30_38 | ~ spl30_41 | ~ spl30_121)), inference(subsumption_resolution, [], [f2652, f212])).
fof(f2652, plain, ((e0 = e2) | (~ spl30_38 | ~ spl30_41 | ~ spl30_121)), inference(forward_demodulation, [], [f2648, f558])).
fof(f2648, plain, ((e2 = op(e1, e1)) | (~ spl30_38 | ~ spl30_121)), inference(backward_demodulation, [], [f1031, f545])).
fof(f2651, plain, (~ spl30_38 | spl30_120), inference(avatar_contradiction_clause, [], [f2650])).
fof(f2650, plain, ($false | (~ spl30_38 | spl30_120)), inference(subsumption_resolution, [], [f2647, f545])).
fof(f2647, plain, (~ (e1 = op(e1, e2)) | (~ spl30_38 | spl30_120)), inference(backward_demodulation, [], [f1026, f545])).
fof(f2622, plain, (~ spl30_54 | ~ spl30_59 | spl30_91), inference(avatar_split_clause, [], [f2619, f862, f632, f611])).
fof(f2619, plain, (~ (e1 = op(e0, e2)) | (~ spl30_59 | spl30_91)), inference(backward_demodulation, [], [f863, f634])).
fof(f2596, plain, (~ spl30_15 | ~ spl30_154 | spl30_158), inference(avatar_contradiction_clause, [], [f2595])).
fof(f2595, plain, ($false | (~ spl30_15 | ~ spl30_154 | spl30_158)), inference(subsumption_resolution, [], [f2592, f447])).
fof(f2592, plain, (~ (e2 = op(e3, e0)) | (~ spl30_154 | spl30_158)), inference(backward_demodulation, [], [f1222, f1201])).
fof(f1201, plain, ((e3 = op(op(e0, e0), e0)) | ~ spl30_154), inference(avatar_component_clause, [], [f1200])).
fof(f2586, plain, (~ spl30_58 | ~ spl30_10), inference(avatar_split_clause, [], [f2585, f424, f628])).
fof(f2585, plain, (~ (e1 = op(e0, e1)) | ~ spl30_10), inference(forward_demodulation, [], [f171, f426])).
fof(f2570, plain, (~ spl30_44 | ~ spl30_10 | spl30_100), inference(avatar_split_clause, [], [f2315, f923, f424, f568])).
fof(f2315, plain, (~ (e3 = op(e1, e1)) | (~ spl30_10 | spl30_100)), inference(forward_demodulation, [], [f925, f426])).
fof(f2563, plain, (~ spl30_40 | ~ spl30_35 | spl30_81), inference(avatar_split_clause, [], [f2545, f796, f530, f551])).
fof(f2545, plain, (~ (e3 = op(e1, e2)) | (~ spl30_35 | spl30_81)), inference(backward_demodulation, [], [f797, f532])).
fof(f2543, plain, (~ spl30_37 | ~ spl30_57 | ~ spl30_121), inference(avatar_contradiction_clause, [], [f2542])).
fof(f2542, plain, ($false | (~ spl30_37 | ~ spl30_57 | ~ spl30_121)), inference(subsumption_resolution, [], [f2541, f212])).
fof(f2541, plain, ((e0 = e2) | (~ spl30_37 | ~ spl30_57 | ~ spl30_121)), inference(forward_demodulation, [], [f2535, f626])).
fof(f2535, plain, ((e2 = op(e0, e1)) | (~ spl30_37 | ~ spl30_121)), inference(backward_demodulation, [], [f1031, f541])).
fof(f2540, plain, (~ spl30_37 | ~ spl30_54 | spl30_120), inference(avatar_contradiction_clause, [], [f2539])).
fof(f2539, plain, ($false | (~ spl30_37 | ~ spl30_54 | spl30_120)), inference(subsumption_resolution, [], [f2534, f613])).
fof(f2534, plain, (~ (e1 = op(e0, e2)) | (~ spl30_37 | spl30_120)), inference(backward_demodulation, [], [f1026, f541])).
fof(f2537, plain, (~ spl30_33 | ~ spl30_37), inference(avatar_split_clause, [], [f2529, f539, f522])).
fof(f2529, plain, (~ (e0 = op(e1, e3)) | ~ spl30_37), inference(backward_demodulation, [], [f198, f541])).
fof(f2527, plain, (~ spl30_46 | ~ spl30_48), inference(avatar_contradiction_clause, [], [f2526])).
fof(f2526, plain, ($false | (~ spl30_46 | ~ spl30_48)), inference(subsumption_resolution, [], [f2525, f215])).
fof(f2525, plain, ((e1 = e3) | (~ spl30_46 | ~ spl30_48)), inference(forward_demodulation, [], [f587, f579])).
fof(f2500, plain, (~ spl30_40 | ~ spl30_18 | spl30_107), inference(avatar_split_clause, [], [f2499, f958, f458, f551])).
fof(f2499, plain, (~ (e3 = op(e1, e2)) | (~ spl30_18 | spl30_107)), inference(forward_demodulation, [], [f959, f460])).
fof(f959, plain, (~ (e3 = op(op(e2, e3), e2)) | spl30_107), inference(avatar_component_clause, [], [f958])).
fof(f2491, plain, (~ spl30_33 | ~ spl30_1), inference(avatar_split_clause, [], [f2490, f386, f522])).
fof(f2490, plain, (~ (e0 = op(e1, e3)) | ~ spl30_1), inference(forward_demodulation, [], [f185, f388])).
fof(f2484, plain, (~ spl30_35 | ~ spl30_18 | spl30_106), inference(avatar_split_clause, [], [f2313, f953, f458, f530])).
fof(f2313, plain, (~ (e2 = op(e1, e3)) | (~ spl30_18 | spl30_106)), inference(forward_demodulation, [], [f955, f460])).
fof(f2477, plain, (~ spl30_6 | ~ spl30_10), inference(avatar_split_clause, [], [f2476, f424, f407])).
fof(f2476, plain, (~ (e1 = op(e3, e2)) | ~ spl30_10), inference(forward_demodulation, [], [f208, f426])).
fof(f2464, plain, (~ spl30_10 | ~ spl30_12), inference(avatar_contradiction_clause, [], [f2463])).
fof(f2463, plain, ($false | (~ spl30_10 | ~ spl30_12)), inference(subsumption_resolution, [], [f2462, f215])).
fof(f2462, plain, ((e1 = e3) | (~ spl30_10 | ~ spl30_12)), inference(forward_demodulation, [], [f434, f426])).
fof(f2446, plain, (~ spl30_18 | ~ spl30_20), inference(avatar_contradiction_clause, [], [f2445])).
fof(f2445, plain, ($false | (~ spl30_18 | ~ spl30_20)), inference(subsumption_resolution, [], [f2444, f215])).
fof(f2444, plain, ((e1 = e3) | (~ spl30_18 | ~ spl30_20)), inference(forward_demodulation, [], [f468, f460])).
fof(f2424, plain, (~ spl30_15 | ~ spl30_64 | spl30_153), inference(avatar_split_clause, [], [f2418, f1194, f653, f445])).
fof(f2418, plain, (~ (e2 = op(e3, e0)) | (~ spl30_64 | spl30_153)), inference(backward_demodulation, [], [f1196, f655])).
fof(f1196, plain, (~ (e2 = op(op(e0, e0), e0)) | spl30_153), inference(avatar_component_clause, [], [f1194])).
fof(f2399, plain, (spl30_28 | ~ spl30_18 | ~ spl30_73), inference(avatar_split_clause, [], [f2257, f740, f458, f500])).
fof(f2257, plain, ((e3 = op(e2, e1)) | (~ spl30_18 | ~ spl30_73)), inference(backward_demodulation, [], [f742, f460])).
fof(f2383, plain, (~ spl30_10 | ~ spl30_40 | ~ spl30_121), inference(avatar_contradiction_clause, [], [f2382])).
fof(f2382, plain, ($false | (~ spl30_10 | ~ spl30_40 | ~ spl30_121)), inference(subsumption_resolution, [], [f2381, f214])).
fof(f2381, plain, ((e1 = e2) | (~ spl30_10 | ~ spl30_40 | ~ spl30_121)), inference(forward_demodulation, [], [f2380, f426])).
fof(f2380, plain, ((e2 = op(e3, e1)) | (~ spl30_40 | ~ spl30_121)), inference(forward_demodulation, [], [f1031, f553])).
fof(f2376, plain, (~ spl30_3 | ~ spl30_18 | spl30_142), inference(avatar_contradiction_clause, [], [f2375])).
fof(f2375, plain, ($false | (~ spl30_3 | ~ spl30_18 | spl30_142)), inference(subsumption_resolution, [], [f2374, f460])).
fof(f2374, plain, (~ (e1 = op(e2, e3)) | (~ spl30_3 | spl30_142)), inference(forward_demodulation, [], [f1145, f396])).
fof(f2366, plain, (~ spl30_46 | spl30_125), inference(avatar_contradiction_clause, [], [f2365])).
fof(f2365, plain, ($false | (~ spl30_46 | spl30_125)), inference(subsumption_resolution, [], [f2363, f579])).
fof(f2363, plain, (~ (e1 = op(e1, e0)) | (~ spl30_46 | spl30_125)), inference(backward_demodulation, [], [f1052, f579])).
fof(f2337, plain, (spl30_53 | ~ spl30_63 | ~ spl30_92), inference(avatar_contradiction_clause, [], [f2336])).
fof(f2336, plain, ($false | (spl30_53 | ~ spl30_63 | ~ spl30_92)), inference(subsumption_resolution, [], [f2329, f608])).
fof(f608, plain, (~ (e0 = op(e0, e2)) | spl30_53), inference(avatar_component_clause, [], [f607])).
fof(f2329, plain, ((e0 = op(e0, e2)) | (~ spl30_63 | ~ spl30_92)), inference(backward_demodulation, [], [f869, f651])).
fof(f2320, plain, (~ spl30_62 | ~ spl30_57 | spl30_91), inference(avatar_split_clause, [], [f2319, f862, f624, f645])).
fof(f2319, plain, (~ (op(e0, e0) = e1) | (~ spl30_57 | spl30_91)), inference(forward_demodulation, [], [f863, f626])).
fof(f2318, plain, (~ spl30_56 | ~ spl30_5 | spl30_97), inference(avatar_split_clause, [], [f2317, f908, f403, f619])).
fof(f2317, plain, (~ (e3 = op(e0, e2)) | (~ spl30_5 | spl30_97)), inference(forward_demodulation, [], [f910, f405])).
fof(f2312, plain, (~ spl30_50 | ~ spl30_33 | spl30_117), inference(avatar_split_clause, [], [f2311, f1009, f522, f594])).
fof(f2311, plain, (~ (e1 = op(e0, e3)) | (~ spl30_33 | spl30_117)), inference(forward_demodulation, [], [f1011, f524])).
fof(f2297, plain, (~ spl30_3 | ~ spl30_18 | ~ spl30_33 | spl30_141), inference(avatar_contradiction_clause, [], [f2296])).
fof(f2296, plain, ($false | (~ spl30_3 | ~ spl30_18 | ~ spl30_33 | spl30_141)), inference(subsumption_resolution, [], [f2295, f524])).
fof(f2295, plain, (~ (e0 = op(e1, e3)) | (~ spl30_3 | ~ spl30_18 | spl30_141)), inference(forward_demodulation, [], [f2294, f460])).
fof(f2294, plain, (~ (e0 = op(op(e2, e3), e3)) | (~ spl30_3 | spl30_141)), inference(forward_demodulation, [], [f1141, f396])).
fof(f1141, plain, (~ (e0 = op(op(op(e3, e3), e3), e3)) | spl30_141), inference(avatar_component_clause, [], [f1139])).
fof(f1139, plain, (spl30_141 <=> (e0 = op(op(op(e3, e3), e3), e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_141])])).
fof(f2284, plain, (spl30_40 | ~ spl30_18 | ~ spl30_107), inference(avatar_split_clause, [], [f2259, f958, f458, f551])).
fof(f2259, plain, ((e3 = op(e1, e2)) | (~ spl30_18 | ~ spl30_107)), inference(backward_demodulation, [], [f960, f460])).
fof(f2283, plain, (~ spl30_3 | spl30_17 | ~ spl30_151), inference(avatar_contradiction_clause, [], [f2282])).
fof(f2282, plain, ($false | (~ spl30_3 | spl30_17 | ~ spl30_151)), inference(subsumption_resolution, [], [f2281, f455])).
fof(f455, plain, (~ (e0 = op(e2, e3)) | spl30_17), inference(avatar_component_clause, [], [f454])).
fof(f2281, plain, ((e0 = op(e2, e3)) | (~ spl30_3 | ~ spl30_151)), inference(backward_demodulation, [], [f1186, f396])).
fof(f2278, plain, (~ spl30_10 | ~ spl30_43 | ~ spl30_100), inference(avatar_contradiction_clause, [], [f2277])).
fof(f2277, plain, ($false | (~ spl30_10 | ~ spl30_43 | ~ spl30_100)), inference(subsumption_resolution, [], [f2276, f216])).
fof(f2276, plain, ((e2 = e3) | (~ spl30_10 | ~ spl30_43 | ~ spl30_100)), inference(forward_demodulation, [], [f2271, f566])).
fof(f2271, plain, ((e3 = op(e1, e1)) | (~ spl30_10 | ~ spl30_100)), inference(backward_demodulation, [], [f924, f426])).
fof(f2275, plain, (~ spl30_10 | spl30_68), inference(avatar_contradiction_clause, [], [f2274])).
fof(f2274, plain, ($false | (~ spl30_10 | spl30_68)), inference(subsumption_resolution, [], [f2270, f426])).
fof(f2270, plain, (~ (e1 = op(e3, e1)) | (~ spl30_10 | spl30_68)), inference(backward_demodulation, [], [f705, f426])).
fof(f2266, plain, (~ spl30_18 | ~ spl30_38 | ~ spl30_107), inference(avatar_contradiction_clause, [], [f2265])).
fof(f2265, plain, ($false | (~ spl30_18 | ~ spl30_38 | ~ spl30_107)), inference(subsumption_resolution, [], [f2264, f215])).
fof(f2264, plain, ((e1 = e3) | (~ spl30_18 | ~ spl30_38 | ~ spl30_107)), inference(forward_demodulation, [], [f2259, f545])).
fof(f2241, plain, (~ spl30_33 | ~ spl30_36), inference(avatar_contradiction_clause, [], [f2240])).
fof(f2240, plain, ($false | (~ spl30_33 | ~ spl30_36)), inference(subsumption_resolution, [], [f2239, f213])).
fof(f2239, plain, ((e0 = e3) | (~ spl30_33 | ~ spl30_36)), inference(forward_demodulation, [], [f536, f524])).
fof(f2228, plain, (~ spl30_39 | ~ spl30_43), inference(avatar_split_clause, [], [f2221, f564, f547])).
fof(f2221, plain, (~ (e2 = op(e1, e2)) | ~ spl30_43), inference(backward_demodulation, [], [f196, f566])).
fof(f2194, plain, (~ spl30_57 | ~ spl30_59), inference(avatar_contradiction_clause, [], [f2193])).
fof(f2193, plain, ($false | (~ spl30_57 | ~ spl30_59)), inference(subsumption_resolution, [], [f2192, f212])).
fof(f2192, plain, ((e0 = e2) | (~ spl30_57 | ~ spl30_59)), inference(forward_demodulation, [], [f634, f626])).
fof(f2185, plain, (~ spl30_48 | ~ spl30_62 | spl30_154), inference(avatar_split_clause, [], [f2184, f1200, f645, f585])).
fof(f2184, plain, (~ (e3 = op(e1, e0)) | (~ spl30_62 | spl30_154)), inference(forward_demodulation, [], [f1202, f647])).
fof(f2182, plain, (~ spl30_50 | ~ spl30_62), inference(avatar_split_clause, [], [f2102, f645, f594])).
fof(f2102, plain, (~ (e1 = op(e0, e3)) | ~ spl30_62), inference(forward_demodulation, [], [f189, f647])).
fof(f2178, plain, (~ spl30_47 | ~ spl30_62 | spl30_153), inference(avatar_split_clause, [], [f2003, f1194, f645, f581])).
fof(f2003, plain, (~ (e2 = op(e1, e0)) | (~ spl30_62 | spl30_153)), inference(backward_demodulation, [], [f1196, f647])).
fof(f2160, plain, (~ spl30_50 | spl30_150 | ~ spl30_151), inference(avatar_split_clause, [], [f2159, f1185, f1181, f594])).
fof(f2159, plain, (~ (e1 = op(e0, e3)) | (spl30_150 | ~ spl30_151)), inference(backward_demodulation, [], [f1183, f1186])).
fof(f2150, plain, (~ spl30_17 | ~ spl30_29), inference(avatar_split_clause, [], [f2146, f505, f454])).
fof(f2146, plain, (~ (e0 = op(e2, e3)) | ~ spl30_29), inference(backward_demodulation, [], [f201, f507])).
fof(f2149, plain, (~ spl30_29 | spl30_76), inference(avatar_contradiction_clause, [], [f2148])).
fof(f2148, plain, ($false | (~ spl30_29 | spl30_76)), inference(subsumption_resolution, [], [f2145, f507])).
fof(f2145, plain, (~ (e0 = op(e2, e0)) | (~ spl30_29 | spl30_76)), inference(backward_demodulation, [], [f756, f507])).
fof(f2132, plain, (spl30_33 | ~ spl30_48 | ~ spl30_84), inference(avatar_split_clause, [], [f2128, f811, f585, f522])).
fof(f2128, plain, ((e0 = op(e1, e3)) | (~ spl30_48 | ~ spl30_84)), inference(backward_demodulation, [], [f813, f587])).
fof(f2131, plain, (~ spl30_36 | ~ spl30_48), inference(avatar_split_clause, [], [f2127, f585, f534])).
fof(f2127, plain, (~ (e3 = op(e1, e3)) | ~ spl30_48), inference(backward_demodulation, [], [f195, f587])).
fof(f2125, plain, (~ spl30_52 | spl30_89), inference(avatar_contradiction_clause, [], [f2124])).
fof(f2124, plain, ($false | (~ spl30_52 | spl30_89)), inference(subsumption_resolution, [], [f2120, f604])).
fof(f2120, plain, (~ (e3 = op(e0, e3)) | (~ spl30_52 | spl30_89)), inference(backward_demodulation, [], [f853, f604])).
fof(f2123, plain, (~ spl30_36 | ~ spl30_52), inference(avatar_split_clause, [], [f2118, f602, f534])).
fof(f2118, plain, (~ (e3 = op(e1, e3)) | ~ spl30_52), inference(backward_demodulation, [], [f181, f604])).
fof(f2117, plain, (~ spl30_62 | ~ spl30_63), inference(avatar_contradiction_clause, [], [f2116])).
fof(f2116, plain, ($false | (~ spl30_62 | ~ spl30_63)), inference(subsumption_resolution, [], [f2115, f214])).
fof(f2115, plain, ((e1 = e2) | (~ spl30_62 | ~ spl30_63)), inference(forward_demodulation, [], [f651, f647])).
fof(f2112, plain, (spl30_33 | ~ spl30_2 | ~ spl30_151), inference(avatar_split_clause, [], [f2111, f1185, f390, f522])).
fof(f2111, plain, ((e0 = op(e1, e3)) | (~ spl30_2 | ~ spl30_151)), inference(forward_demodulation, [], [f1186, f392])).
fof(f2099, plain, (~ spl30_46 | ~ spl30_62), inference(avatar_split_clause, [], [f2098, f645, f577])).
fof(f2098, plain, (~ (e1 = op(e1, e0)) | ~ spl30_62), inference(forward_demodulation, [], [f163, f647])).
fof(f2087, plain, (~ spl30_18 | ~ spl30_22), inference(avatar_split_clause, [], [f2086, f475, f458])).
fof(f2086, plain, (~ (e1 = op(e2, e3)) | ~ spl30_22), inference(forward_demodulation, [], [f204, f477])).
fof(f2084, plain, (~ spl30_18 | ~ spl30_28 | spl30_75), inference(avatar_split_clause, [], [f2055, f750, f500, f458])).
fof(f2055, plain, (~ (e1 = op(e2, e3)) | (~ spl30_28 | spl30_75)), inference(backward_demodulation, [], [f751, f502])).
fof(f2078, plain, (~ spl30_17 | ~ spl30_51 | spl30_128), inference(avatar_contradiction_clause, [], [f2077])).
fof(f2077, plain, ($false | (~ spl30_17 | ~ spl30_51 | spl30_128)), inference(subsumption_resolution, [], [f2076, f456])).
fof(f2076, plain, (~ (e0 = op(e2, e3)) | (~ spl30_51 | spl30_128)), inference(forward_demodulation, [], [f1067, f600])).
fof(f2075, plain, (~ spl30_31 | ~ spl30_51 | ~ spl30_129), inference(avatar_contradiction_clause, [], [f2074])).
fof(f2074, plain, ($false | (~ spl30_31 | ~ spl30_51 | ~ spl30_129)), inference(subsumption_resolution, [], [f2073, f216])).
fof(f2073, plain, ((e2 = e3) | (~ spl30_31 | ~ spl30_51 | ~ spl30_129)), inference(forward_demodulation, [], [f2072, f515])).
fof(f2072, plain, ((e3 = op(e2, e0)) | (~ spl30_51 | ~ spl30_129)), inference(forward_demodulation, [], [f1072, f600])).
fof(f2058, plain, (~ spl30_24 | ~ spl30_28), inference(avatar_split_clause, [], [f2056, f500, f483])).
fof(f2056, plain, (~ (e3 = op(e2, e2)) | ~ spl30_28), inference(backward_demodulation, [], [f202, f502])).
fof(f2026, plain, (~ spl30_28 | ~ spl30_11 | spl30_100), inference(avatar_split_clause, [], [f1991, f923, f428, f500])).
fof(f1991, plain, (~ (e3 = op(e2, e1)) | (~ spl30_11 | spl30_100)), inference(forward_demodulation, [], [f925, f430])).
fof(f2024, plain, (~ spl30_24 | ~ spl30_56), inference(avatar_split_clause, [], [f2023, f619, f483])).
fof(f2023, plain, (~ (e3 = op(e2, e2)) | ~ spl30_56), inference(forward_demodulation, [], [f176, f621])).
fof(f2020, plain, (~ spl30_38 | ~ spl30_39), inference(avatar_contradiction_clause, [], [f2019])).
fof(f2019, plain, ($false | (~ spl30_38 | ~ spl30_39)), inference(subsumption_resolution, [], [f2018, f214])).
fof(f2018, plain, ((e1 = e2) | (~ spl30_38 | ~ spl30_39)), inference(backward_demodulation, [], [f549, f545])).
fof(f2007, plain, (~ spl30_30 | ~ spl30_62), inference(avatar_split_clause, [], [f1998, f645, f509])).
fof(f1998, plain, (~ (e1 = op(e2, e0)) | ~ spl30_62), inference(backward_demodulation, [], [f164, f647])).
fof(f1994, plain, (spl30_56 | ~ spl30_5 | ~ spl30_97), inference(avatar_split_clause, [], [f1993, f908, f403, f619])).
fof(f1993, plain, ((e3 = op(e0, e2)) | (~ spl30_5 | ~ spl30_97)), inference(forward_demodulation, [], [f909, f405])).
fof(f1986, plain, (spl30_62 | ~ spl30_57 | ~ spl30_91), inference(avatar_split_clause, [], [f1872, f862, f624, f645])).
fof(f1872, plain, ((op(e0, e0) = e1) | (~ spl30_57 | ~ spl30_91)), inference(forward_demodulation, [], [f864, f626])).
fof(f1969, plain, (~ spl30_45 | spl30_84), inference(avatar_contradiction_clause, [], [f1968])).
fof(f1968, plain, ($false | (~ spl30_45 | spl30_84)), inference(subsumption_resolution, [], [f1967, f575])).
fof(f1967, plain, (~ (e0 = op(e1, e0)) | (~ spl30_45 | spl30_84)), inference(forward_demodulation, [], [f812, f575])).
fof(f1949, plain, (spl30_34 | ~ spl30_44 | ~ spl30_83), inference(avatar_contradiction_clause, [], [f1948])).
fof(f1948, plain, ($false | (spl30_34 | ~ spl30_44 | ~ spl30_83)), inference(subsumption_resolution, [], [f1947, f527])).
fof(f527, plain, (~ (e1 = op(e1, e3)) | spl30_34), inference(avatar_component_clause, [], [f526])).
fof(f1947, plain, ((e1 = op(e1, e3)) | (~ spl30_44 | ~ spl30_83)), inference(forward_demodulation, [], [f808, f570])).
fof(f1940, plain, (~ spl30_17 | ~ spl30_24 | ~ spl30_74), inference(avatar_contradiction_clause, [], [f1939])).
fof(f1939, plain, ($false | (~ spl30_17 | ~ spl30_24 | ~ spl30_74)), inference(subsumption_resolution, [], [f1938, f212])).
fof(f1938, plain, ((e0 = e2) | (~ spl30_17 | ~ spl30_24 | ~ spl30_74)), inference(forward_demodulation, [], [f1937, f456])).
fof(f1937, plain, ((e2 = op(e2, e3)) | (~ spl30_24 | ~ spl30_74)), inference(forward_demodulation, [], [f747, f485])).
fof(f1932, plain, (~ spl30_54 | ~ spl30_5 | ~ spl30_24 | spl30_148), inference(avatar_split_clause, [], [f1931, f1172, f483, f403, f611])).
fof(f1931, plain, (~ (e1 = op(e0, e2)) | (~ spl30_5 | ~ spl30_24 | spl30_148)), inference(forward_demodulation, [], [f1930, f405])).
fof(f1930, plain, (~ (e1 = op(op(e3, e2), e2)) | (~ spl30_24 | spl30_148)), inference(forward_demodulation, [], [f1174, f485])).
fof(f1913, plain, (~ spl30_23 | ~ spl30_24), inference(avatar_contradiction_clause, [], [f1912])).
fof(f1912, plain, ($false | (~ spl30_23 | ~ spl30_24)), inference(subsumption_resolution, [], [f1911, f216])).
fof(f1911, plain, ((e2 = e3) | (~ spl30_23 | ~ spl30_24)), inference(forward_demodulation, [], [f485, f481])).
fof(f1906, plain, (~ spl30_42 | ~ spl30_44), inference(avatar_contradiction_clause, [], [f1905])).
fof(f1905, plain, ($false | (~ spl30_42 | ~ spl30_44)), inference(subsumption_resolution, [], [f1904, f215])).
fof(f1904, plain, ((e1 = e3) | (~ spl30_42 | ~ spl30_44)), inference(forward_demodulation, [], [f570, f562])).
fof(f1903, plain, (~ spl30_45 | ~ spl30_47), inference(avatar_contradiction_clause, [], [f1902])).
fof(f1902, plain, ($false | (~ spl30_45 | ~ spl30_47)), inference(subsumption_resolution, [], [f1901, f212])).
fof(f1901, plain, ((e0 = e2) | (~ spl30_45 | ~ spl30_47)), inference(backward_demodulation, [], [f583, f575])).
fof(f1878, plain, (~ spl30_54 | ~ spl30_57 | ~ spl30_90), inference(avatar_contradiction_clause, [], [f1877])).
fof(f1877, plain, ($false | (~ spl30_54 | ~ spl30_57 | ~ spl30_90)), inference(subsumption_resolution, [], [f1876, f212])).
fof(f1876, plain, ((e0 = e2) | (~ spl30_54 | ~ spl30_57 | ~ spl30_90)), inference(forward_demodulation, [], [f1875, f626])).
fof(f1875, plain, ((e2 = op(e0, e1)) | (~ spl30_54 | ~ spl30_90)), inference(forward_demodulation, [], [f859, f613])).
fof(f1852, plain, (~ spl30_26 | spl30_75), inference(avatar_contradiction_clause, [], [f1851])).
fof(f1851, plain, ($false | (~ spl30_26 | spl30_75)), inference(subsumption_resolution, [], [f1850, f494])).
fof(f1850, plain, (~ (e1 = op(e2, e1)) | (~ spl30_26 | spl30_75)), inference(forward_demodulation, [], [f751, f494])).
fof(f1839, plain, (~ spl30_13 | ~ spl30_16), inference(avatar_contradiction_clause, [], [f1838])).
fof(f1838, plain, ($false | (~ spl30_13 | ~ spl30_16)), inference(subsumption_resolution, [], [f1836, f213])).
fof(f1836, plain, ((e0 = e3) | (~ spl30_13 | ~ spl30_16)), inference(backward_demodulation, [], [f451, f439])).
fof(f1829, plain, (~ spl30_26 | ~ spl30_28), inference(avatar_contradiction_clause, [], [f1828])).
fof(f1828, plain, ($false | (~ spl30_26 | ~ spl30_28)), inference(subsumption_resolution, [], [f1827, f215])).
fof(f1827, plain, ((e1 = e3) | (~ spl30_26 | ~ spl30_28)), inference(backward_demodulation, [], [f502, f494])).
fof(f1816, plain, (~ spl30_37 | ~ spl30_38), inference(avatar_contradiction_clause, [], [f1815])).
fof(f1815, plain, ($false | (~ spl30_37 | ~ spl30_38)), inference(subsumption_resolution, [], [f1814, f211])).
fof(f1814, plain, ((e0 = e1) | (~ spl30_37 | ~ spl30_38)), inference(backward_demodulation, [], [f545, f541])).
fof(f1813, plain, (~ spl30_6 | ~ spl30_38), inference(avatar_split_clause, [], [f1808, f543, f407])).
fof(f1808, plain, (~ (e1 = op(e3, e2)) | ~ spl30_38), inference(backward_demodulation, [], [f179, f545])).
fof(f1805, plain, (~ spl30_54 | ~ spl30_56), inference(avatar_contradiction_clause, [], [f1804])).
fof(f1804, plain, ($false | (~ spl30_54 | ~ spl30_56)), inference(subsumption_resolution, [], [f1803, f215])).
fof(f1803, plain, ((e1 = e3) | (~ spl30_54 | ~ spl30_56)), inference(forward_demodulation, [], [f621, f613])).
fof(f1802, plain, (~ spl30_57 | spl30_134), inference(avatar_contradiction_clause, [], [f1801])).
fof(f1801, plain, ($false | (~ spl30_57 | spl30_134)), inference(subsumption_resolution, [], [f1798, f626])).
fof(f1798, plain, (~ (e0 = op(e0, e1)) | (~ spl30_57 | spl30_134)), inference(backward_demodulation, [], [f1097, f626])).
fof(f1739, plain, (~ spl30_5 | ~ spl30_49 | ~ spl30_98), inference(avatar_contradiction_clause, [], [f1738])).
fof(f1738, plain, ($false | (~ spl30_5 | ~ spl30_49 | ~ spl30_98)), inference(subsumption_resolution, [], [f1737, f212])).
fof(f1737, plain, ((e0 = e2) | (~ spl30_5 | ~ spl30_49 | ~ spl30_98)), inference(forward_demodulation, [], [f1735, f592])).
fof(f1735, plain, ((e2 = op(e0, e3)) | (~ spl30_5 | ~ spl30_98)), inference(backward_demodulation, [], [f915, f405])).
fof(f1726, plain, (~ spl30_12 | ~ spl30_16), inference(avatar_split_clause, [], [f1724, f449, f432])).
fof(f1724, plain, (~ (e3 = op(e3, e1)) | ~ spl30_16), inference(backward_demodulation, [], [f205, f451])).
fof(f1693, plain, (~ spl30_41 | ~ spl30_42), inference(avatar_contradiction_clause, [], [f1692])).
fof(f1692, plain, ($false | (~ spl30_41 | ~ spl30_42)), inference(subsumption_resolution, [], [f1691, f211])).
fof(f1691, plain, ((e0 = e1) | (~ spl30_41 | ~ spl30_42)), inference(forward_demodulation, [], [f562, f558])).
fof(f1689, plain, (~ spl30_32 | ~ spl30_47 | ~ spl30_62 | spl30_162), inference(avatar_split_clause, [], [f1686, f1242, f645, f581, f517])).
fof(f1686, plain, (~ (e3 = op(e2, e0)) | (~ spl30_47 | ~ spl30_62 | spl30_162)), inference(backward_demodulation, [], [f1611, f583])).
fof(f1611, plain, (~ (e3 = op(op(e1, e0), e0)) | (~ spl30_62 | spl30_162)), inference(backward_demodulation, [], [f1244, f647])).
fof(f1683, plain, (~ spl30_49 | ~ spl30_51), inference(avatar_contradiction_clause, [], [f1682])).
fof(f1682, plain, ($false | (~ spl30_49 | ~ spl30_51)), inference(subsumption_resolution, [], [f1681, f212])).
fof(f1681, plain, ((e0 = e2) | (~ spl30_49 | ~ spl30_51)), inference(forward_demodulation, [], [f600, f592])).
fof(f1672, plain, (~ spl30_11 | ~ spl30_41 | ~ spl30_60 | spl30_155), inference(avatar_split_clause, [], [f1667, f1206, f636, f556, f428])).
fof(f1667, plain, (~ (e2 = op(e3, e1)) | (~ spl30_41 | ~ spl30_60 | spl30_155)), inference(backward_demodulation, [], [f1312, f638])).
fof(f1312, plain, (~ (e2 = op(op(e0, e1), e1)) | (~ spl30_41 | spl30_155)), inference(backward_demodulation, [], [f1208, f558])).
fof(f1634, plain, (~ spl30_21 | ~ spl30_22), inference(avatar_contradiction_clause, [], [f1633])).
fof(f1633, plain, ($false | (~ spl30_21 | ~ spl30_22)), inference(subsumption_resolution, [], [f1632, f211])).
fof(f1632, plain, ((e0 = e1) | (~ spl30_21 | ~ spl30_22)), inference(backward_demodulation, [], [f477, f473])).
fof(f1629, plain, (~ spl30_27 | ~ spl30_28), inference(avatar_contradiction_clause, [], [f1628])).
fof(f1628, plain, ($false | (~ spl30_27 | ~ spl30_28)), inference(subsumption_resolution, [], [f1627, f216])).
fof(f1627, plain, ((e2 = e3) | (~ spl30_27 | ~ spl30_28)), inference(backward_demodulation, [], [f502, f498])).
fof(f1621, plain, (~ spl30_49 | ~ spl30_52), inference(avatar_contradiction_clause, [], [f1620])).
fof(f1620, plain, ($false | (~ spl30_49 | ~ spl30_52)), inference(subsumption_resolution, [], [f1618, f213])).
fof(f1618, plain, ((e0 = e3) | (~ spl30_49 | ~ spl30_52)), inference(backward_demodulation, [], [f604, f592])).
fof(f1613, plain, (~ spl30_14 | ~ spl30_62), inference(avatar_split_clause, [], [f1608, f645, f441])).
fof(f1608, plain, (~ (e1 = op(e3, e0)) | ~ spl30_62), inference(backward_demodulation, [], [f165, f647])).
fof(f1589, plain, (~ spl30_31 | ~ spl30_19), inference(avatar_split_clause, [], [f1584, f462, f513])).
fof(f1584, plain, (~ (e2 = op(e2, e0)) | ~ spl30_19), inference(backward_demodulation, [], [f201, f464])).
fof(f1588, plain, (~ spl30_19 | ~ spl30_22 | ~ spl30_73), inference(avatar_contradiction_clause, [], [f1587])).
fof(f1587, plain, ($false | (~ spl30_19 | ~ spl30_22 | ~ spl30_73)), inference(subsumption_resolution, [], [f1586, f215])).
fof(f1586, plain, ((e1 = e3) | (~ spl30_19 | ~ spl30_22 | ~ spl30_73)), inference(forward_demodulation, [], [f1583, f477])).
fof(f1583, plain, ((e3 = op(e2, e2)) | (~ spl30_19 | ~ spl30_73)), inference(backward_demodulation, [], [f742, f464])).
fof(f1580, plain, (~ spl30_17 | ~ spl30_2 | ~ spl30_35 | spl30_141), inference(avatar_split_clause, [], [f1470, f1139, f530, f390, f454])).
fof(f1470, plain, (~ (e0 = op(e2, e3)) | (~ spl30_2 | ~ spl30_35 | spl30_141)), inference(backward_demodulation, [], [f1426, f532])).
fof(f1426, plain, (~ (e0 = op(op(e1, e3), e3)) | (~ spl30_2 | spl30_141)), inference(backward_demodulation, [], [f1141, f392])).
fof(f1565, plain, (~ spl30_7 | spl30_67), inference(avatar_contradiction_clause, [], [f1564])).
fof(f1564, plain, ($false | (~ spl30_7 | spl30_67)), inference(subsumption_resolution, [], [f1559, f413])).
fof(f1559, plain, (~ (e2 = op(e3, e2)) | (~ spl30_7 | spl30_67)), inference(backward_demodulation, [], [f700, f413])).
fof(f1562, plain, (spl30_19 | ~ spl30_7 | ~ spl30_98), inference(avatar_split_clause, [], [f1557, f913, f411, f462])).
fof(f1557, plain, ((e2 = op(e2, e3)) | (~ spl30_7 | ~ spl30_98)), inference(backward_demodulation, [], [f915, f413])).
fof(f1554, plain, (~ spl30_7 | ~ spl30_55), inference(avatar_split_clause, [], [f1553, f615, f411])).
fof(f1553, plain, (~ (e2 = op(e3, e2)) | ~ spl30_55), inference(forward_demodulation, [], [f177, f617])).
fof(f1552, plain, (~ spl30_5 | ~ spl30_22 | ~ spl30_40 | spl30_139), inference(avatar_split_clause, [], [f1539, f1130, f551, f475, f403])).
fof(f1539, plain, (~ (e0 = op(e3, e2)) | (~ spl30_22 | ~ spl30_40 | spl30_139)), inference(forward_demodulation, [], [f1534, f553])).
fof(f1534, plain, (~ (e0 = op(op(e1, e2), e2)) | (~ spl30_22 | spl30_139)), inference(backward_demodulation, [], [f1132, f477])).
fof(f1538, plain, (spl30_27 | ~ spl30_22 | ~ spl30_74), inference(avatar_split_clause, [], [f1533, f745, f475, f496])).
fof(f1533, plain, ((e2 = op(e2, e1)) | (~ spl30_22 | ~ spl30_74)), inference(backward_demodulation, [], [f747, f477])).
fof(f1526, plain, (spl30_21 | ~ spl30_31 | ~ spl30_76), inference(avatar_split_clause, [], [f1523, f755, f513, f471])).
fof(f1523, plain, ((e0 = op(e2, e2)) | (~ spl30_31 | ~ spl30_76)), inference(backward_demodulation, [], [f757, f515])).
fof(f1489, plain, (~ spl30_50 | ~ spl30_52), inference(avatar_contradiction_clause, [], [f1488])).
fof(f1488, plain, ($false | (~ spl30_50 | ~ spl30_52)), inference(subsumption_resolution, [], [f1487, f215])).
fof(f1487, plain, ((e1 = e3) | (~ spl30_50 | ~ spl30_52)), inference(forward_demodulation, [], [f604, f596])).
fof(f1484, plain, (~ spl30_45 | ~ spl30_41), inference(avatar_split_clause, [], [f1483, f556, f573])).
fof(f1483, plain, (~ (e0 = op(e1, e0)) | ~ spl30_41), inference(forward_demodulation, [], [f193, f558])).
fof(f1453, plain, (~ spl30_17 | ~ spl30_19), inference(avatar_contradiction_clause, [], [f1452])).
fof(f1452, plain, ($false | (~ spl30_17 | ~ spl30_19)), inference(subsumption_resolution, [], [f1451, f212])).
fof(f1451, plain, ((e0 = e2) | (~ spl30_17 | ~ spl30_19)), inference(forward_demodulation, [], [f464, f456])).
fof(f1450, plain, (~ spl30_14 | ~ spl30_16), inference(avatar_contradiction_clause, [], [f1449])).
fof(f1449, plain, ($false | (~ spl30_14 | ~ spl30_16)), inference(subsumption_resolution, [], [f1448, f215])).
fof(f1448, plain, ((e1 = e3) | (~ spl30_14 | ~ spl30_16)), inference(forward_demodulation, [], [f451, f443])).
fof(f1437, plain, (~ spl30_5 | ~ spl30_50 | ~ spl30_98), inference(avatar_contradiction_clause, [], [f1436])).
fof(f1436, plain, ($false | (~ spl30_5 | ~ spl30_50 | ~ spl30_98)), inference(subsumption_resolution, [], [f1435, f214])).
fof(f1435, plain, ((e1 = e2) | (~ spl30_5 | ~ spl30_50 | ~ spl30_98)), inference(forward_demodulation, [], [f1434, f596])).
fof(f1434, plain, ((e2 = op(e0, e3)) | (~ spl30_5 | ~ spl30_98)), inference(forward_demodulation, [], [f915, f405])).
fof(f1433, plain, (~ spl30_5 | ~ spl30_7), inference(avatar_contradiction_clause, [], [f1432])).
fof(f1432, plain, ($false | (~ spl30_5 | ~ spl30_7)), inference(subsumption_resolution, [], [f1431, f212])).
fof(f1431, plain, ((e0 = e2) | (~ spl30_5 | ~ spl30_7)), inference(forward_demodulation, [], [f413, f405])).
fof(f1430, plain, (~ spl30_33 | ~ spl30_2 | spl30_151), inference(avatar_split_clause, [], [f1427, f1185, f390, f522])).
fof(f1427, plain, (~ (e0 = op(e1, e3)) | (~ spl30_2 | spl30_151)), inference(backward_demodulation, [], [f1187, f392])).
fof(f1425, plain, (~ spl30_5 | ~ spl30_6), inference(avatar_contradiction_clause, [], [f1424])).
fof(f1424, plain, ($false | (~ spl30_5 | ~ spl30_6)), inference(subsumption_resolution, [], [f1423, f211])).
fof(f1423, plain, ((e0 = e1) | (~ spl30_5 | ~ spl30_6)), inference(backward_demodulation, [], [f409, f405])).
fof(f1422, plain, (~ spl30_6 | ~ spl30_7), inference(avatar_contradiction_clause, [], [f1421])).
fof(f1421, plain, ($false | (~ spl30_6 | ~ spl30_7)), inference(subsumption_resolution, [], [f1420, f214])).
fof(f1420, plain, ((e1 = e2) | (~ spl30_6 | ~ spl30_7)), inference(backward_demodulation, [], [f413, f409])).
fof(f1419, plain, (~ spl30_7 | ~ spl30_8), inference(avatar_contradiction_clause, [], [f1418])).
fof(f1418, plain, ($false | (~ spl30_7 | ~ spl30_8)), inference(subsumption_resolution, [], [f1417, f216])).
fof(f1417, plain, ((e2 = e3) | (~ spl30_7 | ~ spl30_8)), inference(backward_demodulation, [], [f417, f413])).
fof(f1416, plain, (spl30_3 | ~ spl30_8 | ~ spl30_98), inference(avatar_split_clause, [], [f1411, f913, f415, f394])).
fof(f1411, plain, ((e2 = op(e3, e3)) | (~ spl30_8 | ~ spl30_98)), inference(backward_demodulation, [], [f915, f417])).
fof(f1415, plain, (~ spl30_8 | spl30_97), inference(avatar_contradiction_clause, [], [f1414])).
fof(f1414, plain, ($false | (~ spl30_8 | spl30_97)), inference(subsumption_resolution, [], [f1410, f417])).
fof(f1410, plain, (~ (e3 = op(e3, e2)) | (~ spl30_8 | spl30_97)), inference(backward_demodulation, [], [f910, f417])).
fof(f1407, plain, (~ spl30_10 | ~ spl30_11), inference(avatar_contradiction_clause, [], [f1406])).
fof(f1406, plain, ($false | (~ spl30_10 | ~ spl30_11)), inference(subsumption_resolution, [], [f1405, f214])).
fof(f1405, plain, ((e1 = e2) | (~ spl30_10 | ~ spl30_11)), inference(backward_demodulation, [], [f430, f426])).
fof(f1404, plain, (~ spl30_11 | ~ spl30_12), inference(avatar_contradiction_clause, [], [f1403])).
fof(f1403, plain, ($false | (~ spl30_11 | ~ spl30_12)), inference(subsumption_resolution, [], [f1402, f216])).
fof(f1402, plain, ((e2 = e3) | (~ spl30_11 | ~ spl30_12)), inference(backward_demodulation, [], [f434, f430])).
fof(f1397, plain, (~ spl30_14 | ~ spl30_15), inference(avatar_contradiction_clause, [], [f1396])).
fof(f1396, plain, ($false | (~ spl30_14 | ~ spl30_15)), inference(subsumption_resolution, [], [f1395, f214])).
fof(f1395, plain, ((e1 = e2) | (~ spl30_14 | ~ spl30_15)), inference(backward_demodulation, [], [f447, f443])).
fof(f1394, plain, (~ spl30_15 | ~ spl30_16), inference(avatar_contradiction_clause, [], [f1393])).
fof(f1393, plain, ($false | (~ spl30_15 | ~ spl30_16)), inference(subsumption_resolution, [], [f1392, f216])).
fof(f1392, plain, ((e2 = e3) | (~ spl30_15 | ~ spl30_16)), inference(backward_demodulation, [], [f451, f447])).
fof(f1385, plain, (~ spl30_17 | ~ spl30_18), inference(avatar_contradiction_clause, [], [f1384])).
fof(f1384, plain, ($false | (~ spl30_17 | ~ spl30_18)), inference(subsumption_resolution, [], [f1383, f211])).
fof(f1383, plain, ((e0 = e1) | (~ spl30_17 | ~ spl30_18)), inference(backward_demodulation, [], [f460, f456])).
fof(f1382, plain, (~ spl30_18 | ~ spl30_19), inference(avatar_contradiction_clause, [], [f1381])).
fof(f1381, plain, ($false | (~ spl30_18 | ~ spl30_19)), inference(subsumption_resolution, [], [f1380, f214])).
fof(f1380, plain, ((e1 = e2) | (~ spl30_18 | ~ spl30_19)), inference(backward_demodulation, [], [f464, f460])).
fof(f1379, plain, (~ spl30_19 | ~ spl30_20), inference(avatar_contradiction_clause, [], [f1378])).
fof(f1378, plain, ($false | (~ spl30_19 | ~ spl30_20)), inference(subsumption_resolution, [], [f1377, f216])).
fof(f1377, plain, ((e2 = e3) | (~ spl30_19 | ~ spl30_20)), inference(backward_demodulation, [], [f468, f464])).
fof(f1355, plain, (~ spl30_21 | ~ spl30_25), inference(avatar_split_clause, [], [f1351, f488, f471])).
fof(f1351, plain, (~ (e0 = op(e2, e2)) | ~ spl30_25), inference(backward_demodulation, [], [f202, f490])).
fof(f1340, plain, (~ spl30_34 | ~ spl30_35), inference(avatar_contradiction_clause, [], [f1339])).
fof(f1339, plain, ($false | (~ spl30_34 | ~ spl30_35)), inference(subsumption_resolution, [], [f1338, f214])).
fof(f1338, plain, ((e1 = e2) | (~ spl30_34 | ~ spl30_35)), inference(backward_demodulation, [], [f532, f528])).
fof(f1337, plain, (~ spl30_35 | ~ spl30_36), inference(avatar_contradiction_clause, [], [f1336])).
fof(f1336, plain, ($false | (~ spl30_35 | ~ spl30_36)), inference(subsumption_resolution, [], [f1335, f216])).
fof(f1335, plain, ((e2 = e3) | (~ spl30_35 | ~ spl30_36)), inference(backward_demodulation, [], [f536, f532])).
fof(f1306, plain, (~ spl30_33 | ~ spl30_45), inference(avatar_split_clause, [], [f1301, f573, f522])).
fof(f1301, plain, (~ (e0 = op(e1, e3)) | ~ spl30_45), inference(backward_demodulation, [], [f195, f575])).
fof(f1303, plain, (~ spl30_13 | ~ spl30_45), inference(avatar_split_clause, [], [f1298, f573, f437])).
fof(f1298, plain, (~ (e0 = op(e3, e0)) | ~ spl30_45), inference(backward_demodulation, [], [f167, f575])).
fof(f1296, plain, (~ spl30_50 | ~ spl30_51), inference(avatar_contradiction_clause, [], [f1295])).
fof(f1295, plain, ($false | (~ spl30_50 | ~ spl30_51)), inference(subsumption_resolution, [], [f1294, f214])).
fof(f1294, plain, ((e1 = e2) | (~ spl30_50 | ~ spl30_51)), inference(backward_demodulation, [], [f600, f596])).
fof(f1293, plain, (~ spl30_51 | ~ spl30_52), inference(avatar_contradiction_clause, [], [f1292])).
fof(f1292, plain, ($false | (~ spl30_51 | ~ spl30_52)), inference(subsumption_resolution, [], [f1291, f216])).
fof(f1291, plain, ((e2 = e3) | (~ spl30_51 | ~ spl30_52)), inference(backward_demodulation, [], [f604, f600])).
fof(f1263, plain, (~ spl30_49 | ~ spl30_61), inference(avatar_split_clause, [], [f1254, f641, f590])).
fof(f1254, plain, (~ (e0 = op(e0, e3)) | ~ spl30_61), inference(backward_demodulation, [], [f189, f643])).
fof(f1248, plain, (~ spl30_160 | ~ spl30_144 | ~ spl30_41), inference(avatar_split_clause, [], [f384, f556, f1152, f1232])).
fof(f384, plain, (~ (e0 = op(e1, e1)) | ~ (e2 = op(op(e1, e1), e1)) | ~ (e3 = op(op(op(e1, e1), e1), e1))), inference(cnf_transformation, [], [f53])).
fof(f53, plain, (~ (e0 = op(e1, e1)) | ~ (e2 = op(op(e1, e1), e1)) | ~ (e3 = op(op(op(e1, e1), e1), e1))), inference(ennf_transformation, [], [f29])).
fof(f29, plain, ~ ((e0 = op(e1, e1)) & (e2 = op(op(e1, e1), e1)) & (e3 = op(op(op(e1, e1), e1), e1))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG164+1.p', ax29)).
fof(f1247, plain, (~ spl30_162 | ~ spl30_153 | ~ spl30_62), inference(avatar_split_clause, [], [f383, f645, f1194, f1242])).
fof(f383, plain, (~ (op(e0, e0) = e1) | ~ (e2 = op(op(e0, e0), e0)) | ~ (e3 = op(op(op(e0, e0), e0), e0))), inference(cnf_transformation, [], [f52])).
fof(f52, plain, (~ (op(e0, e0) = e1) | ~ (e2 = op(op(e0, e0), e0)) | ~ (e3 = op(op(op(e0, e0), e0), e0))), inference(ennf_transformation, [], [f28])).
fof(f28, plain, ~ ((op(e0, e0) = e1) & (e2 = op(op(e0, e0), e0)) & (e3 = op(op(op(e0, e0), e0), e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG164+1.p', ax28)).
fof(f1246, plain, (~ spl30_161 | ~ spl30_140 | ~ spl30_21), inference(avatar_split_clause, [], [f382, f471, f1134, f1237])).
fof(f382, plain, (~ (e0 = op(e2, e2)) | ~ (e1 = op(op(e2, e2), e2)) | ~ (e3 = op(op(op(e2, e2), e2), e2))), inference(cnf_transformation, [], [f51])).
fof(f51, plain, (~ (e0 = op(e2, e2)) | ~ (e1 = op(op(e2, e2), e2)) | ~ (e3 = op(op(op(e2, e2), e2), e2))), inference(ennf_transformation, [], [f27])).
fof(f27, plain, ~ ((e0 = op(e2, e2)) & (e1 = op(op(e2, e2), e2)) & (e3 = op(op(op(e2, e2), e2), e2))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG164+1.p', ax27)).
fof(f1245, plain, (~ spl30_162 | ~ spl30_159 | ~ spl30_63), inference(avatar_split_clause, [], [f381, f649, f1224, f1242])).
fof(f381, plain, (~ (op(e0, e0) = e2) | ~ (e1 = op(op(e0, e0), e0)) | ~ (e3 = op(op(op(e0, e0), e0), e0))), inference(cnf_transformation, [], [f50])).
fof(f50, plain, (~ (op(e0, e0) = e2) | ~ (e1 = op(op(e0, e0), e0)) | ~ (e3 = op(op(op(e0, e0), e0), e0))), inference(ennf_transformation, [], [f26])).
fof(f26, plain, ~ ((op(e0, e0) = e2) & (e1 = op(op(e0, e0), e0)) & (e3 = op(op(op(e0, e0), e0), e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG164+1.p', ax26)).
fof(f1240, plain, (~ spl30_161 | ~ spl30_149 | ~ spl30_22), inference(avatar_split_clause, [], [f380, f475, f1176, f1237])).
fof(f380, plain, (~ (e1 = op(e2, e2)) | ~ (e0 = op(op(e2, e2), e2)) | ~ (e3 = op(op(op(e2, e2), e2), e2))), inference(cnf_transformation, [], [f49])).
fof(f49, plain, (~ (e1 = op(e2, e2)) | ~ (e0 = op(op(e2, e2), e2)) | ~ (e3 = op(op(op(e2, e2), e2), e2))), inference(ennf_transformation, [], [f25])).
fof(f25, plain, ~ ((e1 = op(e2, e2)) & (e0 = op(op(e2, e2), e2)) & (e3 = op(op(op(e2, e2), e2), e2))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG164+1.p', ax25)).
fof(f1235, plain, (~ spl30_160 | ~ spl30_156 | ~ spl30_43), inference(avatar_split_clause, [], [f379, f564, f1210, f1232])).
fof(f379, plain, (~ (e2 = op(e1, e1)) | ~ (e0 = op(op(e1, e1), e1)) | ~ (e3 = op(op(op(e1, e1), e1), e1))), inference(cnf_transformation, [], [f48])).
fof(f48, plain, (~ (e2 = op(e1, e1)) | ~ (e0 = op(op(e1, e1), e1)) | ~ (e3 = op(op(op(e1, e1), e1), e1))), inference(ennf_transformation, [], [f24])).
fof(f24, plain, ~ ((e2 = op(e1, e1)) & (e0 = op(op(e1, e1), e1)) & (e3 = op(op(op(e1, e1), e1), e1))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG164+1.p', ax24)).
fof(f1230, plain, (~ spl30_155 | ~ spl30_146 | ~ spl30_41), inference(avatar_split_clause, [], [f378, f556, f1162, f1206])).
fof(f378, plain, (~ (e0 = op(e1, e1)) | ~ (e3 = op(op(e1, e1), e1)) | ~ (e2 = op(op(op(e1, e1), e1), e1))), inference(cnf_transformation, [], [f47])).
fof(f47, plain, (~ (e0 = op(e1, e1)) | ~ (e3 = op(op(e1, e1), e1)) | ~ (e2 = op(op(op(e1, e1), e1), e1))), inference(ennf_transformation, [], [f23])).
fof(f23, plain, ~ ((e0 = op(e1, e1)) & (e3 = op(op(e1, e1), e1)) & (e2 = op(op(op(e1, e1), e1), e1))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG164+1.p', ax23)).
fof(f1229, plain, (~ spl30_158 | ~ spl30_154 | ~ spl30_62), inference(avatar_split_clause, [], [f377, f645, f1200, f1220])).
fof(f377, plain, (~ (op(e0, e0) = e1) | ~ (e3 = op(op(e0, e0), e0)) | ~ (e2 = op(op(op(e0, e0), e0), e0))), inference(cnf_transformation, [], [f46])).
fof(f46, plain, (~ (op(e0, e0) = e1) | ~ (e3 = op(op(e0, e0), e0)) | ~ (e2 = op(op(op(e0, e0), e0), e0))), inference(ennf_transformation, [], [f22])).
fof(f22, plain, ~ ((op(e0, e0) = e1) & (e3 = op(op(e0, e0), e0)) & (e2 = op(op(op(e0, e0), e0), e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG164+1.p', ax22)).
fof(f1228, plain, (~ spl30_157 | ~ spl30_142 | ~ spl30_1), inference(avatar_split_clause, [], [f376, f386, f1143, f1215])).
fof(f376, plain, (~ (e0 = op(e3, e3)) | ~ (e1 = op(op(e3, e3), e3)) | ~ (e2 = op(op(op(e3, e3), e3), e3))), inference(cnf_transformation, [], [f45])).
fof(f45, plain, (~ (e0 = op(e3, e3)) | ~ (e1 = op(op(e3, e3), e3)) | ~ (e2 = op(op(op(e3, e3), e3), e3))), inference(ennf_transformation, [], [f21])).
fof(f21, plain, ~ ((e0 = op(e3, e3)) & (e1 = op(op(e3, e3), e3)) & (e2 = op(op(op(e3, e3), e3), e3))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG164+1.p', ax21)).
fof(f1227, plain, (~ spl30_158 | ~ spl30_159 | ~ spl30_64), inference(avatar_split_clause, [], [f375, f653, f1224, f1220])).
fof(f375, plain, (~ (op(e0, e0) = e3) | ~ (e1 = op(op(e0, e0), e0)) | ~ (e2 = op(op(op(e0, e0), e0), e0))), inference(cnf_transformation, [], [f44])).
fof(f44, plain, (~ (op(e0, e0) = e3) | ~ (e1 = op(op(e0, e0), e0)) | ~ (e2 = op(op(op(e0, e0), e0), e0))), inference(ennf_transformation, [], [f20])).
fof(f20, plain, ~ ((op(e0, e0) = e3) & (e1 = op(op(e0, e0), e0)) & (e2 = op(op(op(e0, e0), e0), e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG164+1.p', ax20)).
fof(f1218, plain, (~ spl30_157 | ~ spl30_151 | ~ spl30_2), inference(avatar_split_clause, [], [f374, f390, f1185, f1215])).
fof(f374, plain, (~ (e1 = op(e3, e3)) | ~ (e0 = op(op(e3, e3), e3)) | ~ (e2 = op(op(op(e3, e3), e3), e3))), inference(cnf_transformation, [], [f43])).
fof(f43, plain, (~ (e1 = op(e3, e3)) | ~ (e0 = op(op(e3, e3), e3)) | ~ (e2 = op(op(op(e3, e3), e3), e3))), inference(ennf_transformation, [], [f19])).
fof(f19, plain, ~ ((e1 = op(e3, e3)) & (e0 = op(op(e3, e3), e3)) & (e2 = op(op(op(e3, e3), e3), e3))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG164+1.p', ax19)).
fof(f1213, plain, (~ spl30_155 | ~ spl30_156 | ~ spl30_44), inference(avatar_split_clause, [], [f373, f568, f1210, f1206])).
fof(f373, plain, (~ (e3 = op(e1, e1)) | ~ (e0 = op(op(e1, e1), e1)) | ~ (e2 = op(op(op(e1, e1), e1), e1))), inference(cnf_transformation, [], [f42])).
fof(f42, plain, (~ (e3 = op(e1, e1)) | ~ (e0 = op(op(e1, e1), e1)) | ~ (e2 = op(op(op(e1, e1), e1), e1))), inference(ennf_transformation, [], [f18])).
fof(f18, plain, ~ ((e3 = op(e1, e1)) & (e0 = op(op(e1, e1), e1)) & (e2 = op(op(op(e1, e1), e1), e1))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG164+1.p', ax18)).
fof(f1204, plain, (~ spl30_148 | ~ spl30_147 | ~ spl30_21), inference(avatar_split_clause, [], [f372, f471, f1167, f1172])).
fof(f372, plain, (~ (e0 = op(e2, e2)) | ~ (e3 = op(op(e2, e2), e2)) | ~ (e1 = op(op(op(e2, e2), e2), e2))), inference(cnf_transformation, [], [f41])).
fof(f41, plain, (~ (e0 = op(e2, e2)) | ~ (e3 = op(op(e2, e2), e2)) | ~ (e1 = op(op(op(e2, e2), e2), e2))), inference(ennf_transformation, [], [f17])).
fof(f17, plain, ~ ((e0 = op(e2, e2)) & (e3 = op(op(e2, e2), e2)) & (e1 = op(op(op(e2, e2), e2), e2))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG164+1.p', ax17)).
fof(f1203, plain, (~ spl30_152 | ~ spl30_154 | ~ spl30_63), inference(avatar_split_clause, [], [f371, f649, f1200, f1190])).
fof(f371, plain, (~ (op(e0, e0) = e2) | ~ (e3 = op(op(e0, e0), e0)) | ~ (e1 = op(op(op(e0, e0), e0), e0))), inference(cnf_transformation, [], [f40])).
fof(f40, plain, (~ (op(e0, e0) = e2) | ~ (e3 = op(op(e0, e0), e0)) | ~ (e1 = op(op(op(e0, e0), e0), e0))), inference(ennf_transformation, [], [f16])).
fof(f16, plain, ~ ((op(e0, e0) = e2) & (e3 = op(op(e0, e0), e0)) & (e1 = op(op(op(e0, e0), e0), e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG164+1.p', ax16)).
fof(f1198, plain, (~ spl30_150 | ~ spl30_145 | ~ spl30_1), inference(avatar_split_clause, [], [f370, f386, f1157, f1181])).
fof(f370, plain, (~ (e0 = op(e3, e3)) | ~ (e2 = op(op(e3, e3), e3)) | ~ (e1 = op(op(op(e3, e3), e3), e3))), inference(cnf_transformation, [], [f39])).
fof(f39, plain, (~ (e0 = op(e3, e3)) | ~ (e2 = op(op(e3, e3), e3)) | ~ (e1 = op(op(op(e3, e3), e3), e3))), inference(ennf_transformation, [], [f15])).
fof(f15, plain, ~ ((e0 = op(e3, e3)) & (e2 = op(op(e3, e3), e3)) & (e1 = op(op(op(e3, e3), e3), e3))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG164+1.p', ax15)).
fof(f1197, plain, (~ spl30_152 | ~ spl30_153 | ~ spl30_64), inference(avatar_split_clause, [], [f369, f653, f1194, f1190])).
fof(f369, plain, (~ (op(e0, e0) = e3) | ~ (e2 = op(op(e0, e0), e0)) | ~ (e1 = op(op(op(e0, e0), e0), e0))), inference(cnf_transformation, [], [f38])).
fof(f38, plain, (~ (op(e0, e0) = e3) | ~ (e2 = op(op(e0, e0), e0)) | ~ (e1 = op(op(op(e0, e0), e0), e0))), inference(ennf_transformation, [], [f14])).
fof(f14, plain, ~ ((op(e0, e0) = e3) & (e2 = op(op(e0, e0), e0)) & (e1 = op(op(op(e0, e0), e0), e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG164+1.p', ax14)).
fof(f1188, plain, (~ spl30_150 | ~ spl30_151 | ~ spl30_3), inference(avatar_split_clause, [], [f368, f394, f1185, f1181])).
fof(f368, plain, (~ (e2 = op(e3, e3)) | ~ (e0 = op(op(e3, e3), e3)) | ~ (e1 = op(op(op(e3, e3), e3), e3))), inference(cnf_transformation, [], [f37])).
fof(f37, plain, (~ (e2 = op(e3, e3)) | ~ (e0 = op(op(e3, e3), e3)) | ~ (e1 = op(op(op(e3, e3), e3), e3))), inference(ennf_transformation, [], [f13])).
fof(f13, plain, ~ ((e2 = op(e3, e3)) & (e0 = op(op(e3, e3), e3)) & (e1 = op(op(op(e3, e3), e3), e3))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG164+1.p', ax13)).
fof(f1179, plain, (~ spl30_148 | ~ spl30_149 | ~ spl30_24), inference(avatar_split_clause, [], [f367, f483, f1176, f1172])).
fof(f367, plain, (~ (e3 = op(e2, e2)) | ~ (e0 = op(op(e2, e2), e2)) | ~ (e1 = op(op(op(e2, e2), e2), e2))), inference(cnf_transformation, [], [f36])).
fof(f36, plain, (~ (e3 = op(e2, e2)) | ~ (e0 = op(op(e2, e2), e2)) | ~ (e1 = op(op(op(e2, e2), e2), e2))), inference(ennf_transformation, [], [f12])).
fof(f12, plain, ~ ((e3 = op(e2, e2)) & (e0 = op(op(e2, e2), e2)) & (e1 = op(op(op(e2, e2), e2), e2))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG164+1.p', ax12)).
fof(f1170, plain, (~ spl30_139 | ~ spl30_147 | ~ spl30_22), inference(avatar_split_clause, [], [f366, f475, f1167, f1130])).
fof(f366, plain, (~ (e1 = op(e2, e2)) | ~ (e3 = op(op(e2, e2), e2)) | ~ (e0 = op(op(op(e2, e2), e2), e2))), inference(cnf_transformation, [], [f35])).
fof(f35, plain, (~ (e1 = op(e2, e2)) | ~ (e3 = op(op(e2, e2), e2)) | ~ (e0 = op(op(op(e2, e2), e2), e2))), inference(ennf_transformation, [], [f11])).
fof(f11, plain, ~ ((e1 = op(e2, e2)) & (e3 = op(op(e2, e2), e2)) & (e0 = op(op(op(e2, e2), e2), e2))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG164+1.p', ax11)).
fof(f1165, plain, (~ spl30_143 | ~ spl30_146 | ~ spl30_43), inference(avatar_split_clause, [], [f365, f564, f1162, f1148])).
fof(f365, plain, (~ (e2 = op(e1, e1)) | ~ (e3 = op(op(e1, e1), e1)) | ~ (e0 = op(op(op(e1, e1), e1), e1))), inference(cnf_transformation, [], [f34])).
fof(f34, plain, (~ (e2 = op(e1, e1)) | ~ (e3 = op(op(e1, e1), e1)) | ~ (e0 = op(op(op(e1, e1), e1), e1))), inference(ennf_transformation, [], [f10])).
fof(f10, plain, ~ ((e2 = op(e1, e1)) & (e3 = op(op(e1, e1), e1)) & (e0 = op(op(op(e1, e1), e1), e1))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG164+1.p', ax10)).
fof(f1160, plain, (~ spl30_141 | ~ spl30_145 | ~ spl30_2), inference(avatar_split_clause, [], [f364, f390, f1157, f1139])).
fof(f364, plain, (~ (e1 = op(e3, e3)) | ~ (e2 = op(op(e3, e3), e3)) | ~ (e0 = op(op(op(e3, e3), e3), e3))), inference(cnf_transformation, [], [f33])).
fof(f33, plain, (~ (e1 = op(e3, e3)) | ~ (e2 = op(op(e3, e3), e3)) | ~ (e0 = op(op(op(e3, e3), e3), e3))), inference(ennf_transformation, [], [f9])).
fof(f9, plain, ~ ((e1 = op(e3, e3)) & (e2 = op(op(e3, e3), e3)) & (e0 = op(op(op(e3, e3), e3), e3))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG164+1.p', ax9)).
fof(f1155, plain, (~ spl30_143 | ~ spl30_144 | ~ spl30_44), inference(avatar_split_clause, [], [f363, f568, f1152, f1148])).
fof(f363, plain, (~ (e3 = op(e1, e1)) | ~ (e2 = op(op(e1, e1), e1)) | ~ (e0 = op(op(op(e1, e1), e1), e1))), inference(cnf_transformation, [], [f32])).
fof(f32, plain, (~ (e3 = op(e1, e1)) | ~ (e2 = op(op(e1, e1), e1)) | ~ (e0 = op(op(op(e1, e1), e1), e1))), inference(ennf_transformation, [], [f8])).
fof(f8, plain, ~ ((e3 = op(e1, e1)) & (e2 = op(op(e1, e1), e1)) & (e0 = op(op(op(e1, e1), e1), e1))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG164+1.p', ax8)).
fof(f1146, plain, (~ spl30_141 | ~ spl30_142 | ~ spl30_3), inference(avatar_split_clause, [], [f362, f394, f1143, f1139])).
fof(f362, plain, (~ (e2 = op(e3, e3)) | ~ (e1 = op(op(e3, e3), e3)) | ~ (e0 = op(op(op(e3, e3), e3), e3))), inference(cnf_transformation, [], [f31])).
fof(f31, plain, (~ (e2 = op(e3, e3)) | ~ (e1 = op(op(e3, e3), e3)) | ~ (e0 = op(op(op(e3, e3), e3), e3))), inference(ennf_transformation, [], [f7])).
fof(f7, plain, ~ ((e2 = op(e3, e3)) & (e1 = op(op(e3, e3), e3)) & (e0 = op(op(op(e3, e3), e3), e3))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG164+1.p', ax7)).
fof(f1137, plain, (~ spl30_139 | ~ spl30_140 | ~ spl30_24), inference(avatar_split_clause, [], [f361, f483, f1134, f1130])).
fof(f361, plain, (~ (e3 = op(e2, e2)) | ~ (e1 = op(op(e2, e2), e2)) | ~ (e0 = op(op(op(e2, e2), e2), e2))), inference(cnf_transformation, [], [f30])).
fof(f30, plain, (~ (e3 = op(e2, e2)) | ~ (e1 = op(op(e2, e2), e2)) | ~ (e0 = op(op(op(e2, e2), e2), e2))), inference(ennf_transformation, [], [f6])).
fof(f6, plain, ~ ((e3 = op(e2, e2)) & (e1 = op(op(e2, e2), e2)) & (e0 = op(op(op(e2, e2), e2), e2))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG164+1.p', ax6)).
fof(f1128, plain, (spl30_136 | spl30_133 | spl30_130 | spl30_127 | spl30_124 | spl30_122 | spl30_119 | spl30_116 | spl30_113 | spl30_110 | spl30_108 | spl30_105 | spl30_102 | spl30_99 | spl30_96 | ~ spl30_66), inference(avatar_split_clause, [], [f352, f694, f904, f919, f934, f949, f964, f975, f990, f1005, f1020, f1035, f1046, f1061, f1076, f1091, f1106])).
fof(f1106, plain, (spl30_136 <=> sP0), introduced(avatar_definition, [new_symbols(naming, [spl30_136])])).
fof(f1091, plain, (spl30_133 <=> sP1), introduced(avatar_definition, [new_symbols(naming, [spl30_133])])).
fof(f1076, plain, (spl30_130 <=> sP2), introduced(avatar_definition, [new_symbols(naming, [spl30_130])])).
fof(f1061, plain, (spl30_127 <=> sP3), introduced(avatar_definition, [new_symbols(naming, [spl30_127])])).
fof(f1046, plain, (spl30_124 <=> sP4), introduced(avatar_definition, [new_symbols(naming, [spl30_124])])).
fof(f1035, plain, (spl30_122 <=> sP5), introduced(avatar_definition, [new_symbols(naming, [spl30_122])])).
fof(f1020, plain, (spl30_119 <=> sP6), introduced(avatar_definition, [new_symbols(naming, [spl30_119])])).
fof(f1005, plain, (spl30_116 <=> sP7), introduced(avatar_definition, [new_symbols(naming, [spl30_116])])).
fof(f990, plain, (spl30_113 <=> sP8), introduced(avatar_definition, [new_symbols(naming, [spl30_113])])).
fof(f975, plain, (spl30_110 <=> sP9), introduced(avatar_definition, [new_symbols(naming, [spl30_110])])).
fof(f964, plain, (spl30_108 <=> sP10), introduced(avatar_definition, [new_symbols(naming, [spl30_108])])).
fof(f949, plain, (spl30_105 <=> sP11), introduced(avatar_definition, [new_symbols(naming, [spl30_105])])).
fof(f934, plain, (spl30_102 <=> sP12), introduced(avatar_definition, [new_symbols(naming, [spl30_102])])).
fof(f919, plain, (spl30_99 <=> sP13), introduced(avatar_definition, [new_symbols(naming, [spl30_99])])).
fof(f904, plain, (spl30_96 <=> sP14), introduced(avatar_definition, [new_symbols(naming, [spl30_96])])).
fof(f352, plain, (~ (e3 = op(e3, op(e3, e3))) | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f84])).
fof(f84, plain, ((((e3 = op(e3, op(e3, e3))) & (e2 = op(e3, op(e3, e2))) & (e1 = op(e3, op(e3, e1))) & (e0 = op(e3, op(e3, e0))) & (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15) & ((~ (e3 = op(op(e3, e3), e3)) & (e3 = op(op(e3, e3), e3)) & ~ (e3 = op(e3, op(e3, e3)))) | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0)), inference(definition_folding, [], [f5, e83, e82, e81, e80, e79, e78, e77, e76, e75, e74, e73, e72, e71, e70, e69, e68, e67, e66, e65, e64, e63, e62, e61, e60, e59, e58, e57, e56, e55, e54])).
fof(f54, plain, ((~ (e0 = op(op(e0, e0), e0)) & (e0 = op(op(e0, e0), e0)) & ~ (e0 = op(e0, op(e0, e0)))) | ~ sP0), inference(usedef, [], [e54])).
fof(e54, plain, (sP0 <=> (~ (e0 = op(op(e0, e0), e0)) & (e0 = op(op(e0, e0), e0)) & ~ (e0 = op(e0, op(e0, e0))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f55, plain, ((~ (e0 = op(op(e0, e1), e1)) & (e1 = op(op(e0, e1), e0)) & ~ (e1 = op(e0, op(e0, e1)))) | ~ sP1), inference(usedef, [], [e55])).
fof(e55, plain, (sP1 <=> (~ (e0 = op(op(e0, e1), e1)) & (e1 = op(op(e0, e1), e0)) & ~ (e1 = op(e0, op(e0, e1))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f56, plain, ((~ (e0 = op(op(e0, e2), e2)) & (e2 = op(op(e0, e2), e0)) & ~ (e2 = op(e0, op(e0, e2)))) | ~ sP2), inference(usedef, [], [e56])).
fof(e56, plain, (sP2 <=> (~ (e0 = op(op(e0, e2), e2)) & (e2 = op(op(e0, e2), e0)) & ~ (e2 = op(e0, op(e0, e2))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f57, plain, ((~ (e0 = op(op(e0, e3), e3)) & (e3 = op(op(e0, e3), e0)) & ~ (e3 = op(e0, op(e0, e3)))) | ~ sP3), inference(usedef, [], [e57])).
fof(e57, plain, (sP3 <=> (~ (e0 = op(op(e0, e3), e3)) & (e3 = op(op(e0, e3), e0)) & ~ (e3 = op(e0, op(e0, e3))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f58, plain, ((~ (e1 = op(op(e1, e0), e0)) & (e0 = op(op(e1, e0), e1)) & ~ (e0 = op(e1, op(e1, e0)))) | ~ sP4), inference(usedef, [], [e58])).
fof(e58, plain, (sP4 <=> (~ (e1 = op(op(e1, e0), e0)) & (e0 = op(op(e1, e0), e1)) & ~ (e0 = op(e1, op(e1, e0))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f59, plain, ((~ (e1 = op(op(e1, e1), e1)) & (e1 = op(op(e1, e1), e1)) & ~ (e1 = op(e1, op(e1, e1)))) | ~ sP5), inference(usedef, [], [e59])).
fof(e59, plain, (sP5 <=> (~ (e1 = op(op(e1, e1), e1)) & (e1 = op(op(e1, e1), e1)) & ~ (e1 = op(e1, op(e1, e1))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f60, plain, ((~ (e1 = op(op(e1, e2), e2)) & (e2 = op(op(e1, e2), e1)) & ~ (e2 = op(e1, op(e1, e2)))) | ~ sP6), inference(usedef, [], [e60])).
fof(e60, plain, (sP6 <=> (~ (e1 = op(op(e1, e2), e2)) & (e2 = op(op(e1, e2), e1)) & ~ (e2 = op(e1, op(e1, e2))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f61, plain, ((~ (e1 = op(op(e1, e3), e3)) & (e3 = op(op(e1, e3), e1)) & ~ (e3 = op(e1, op(e1, e3)))) | ~ sP7), inference(usedef, [], [e61])).
fof(e61, plain, (sP7 <=> (~ (e1 = op(op(e1, e3), e3)) & (e3 = op(op(e1, e3), e1)) & ~ (e3 = op(e1, op(e1, e3))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f62, plain, ((~ (e2 = op(op(e2, e0), e0)) & (e0 = op(op(e2, e0), e2)) & ~ (e0 = op(e2, op(e2, e0)))) | ~ sP8), inference(usedef, [], [e62])).
fof(e62, plain, (sP8 <=> (~ (e2 = op(op(e2, e0), e0)) & (e0 = op(op(e2, e0), e2)) & ~ (e0 = op(e2, op(e2, e0))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f63, plain, ((~ (e2 = op(op(e2, e1), e1)) & (e1 = op(op(e2, e1), e2)) & ~ (e1 = op(e2, op(e2, e1)))) | ~ sP9), inference(usedef, [], [e63])).
fof(e63, plain, (sP9 <=> (~ (e2 = op(op(e2, e1), e1)) & (e1 = op(op(e2, e1), e2)) & ~ (e1 = op(e2, op(e2, e1))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f64, plain, ((~ (e2 = op(op(e2, e2), e2)) & (e2 = op(op(e2, e2), e2)) & ~ (e2 = op(e2, op(e2, e2)))) | ~ sP10), inference(usedef, [], [e64])).
fof(e64, plain, (sP10 <=> (~ (e2 = op(op(e2, e2), e2)) & (e2 = op(op(e2, e2), e2)) & ~ (e2 = op(e2, op(e2, e2))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f65, plain, ((~ (e2 = op(op(e2, e3), e3)) & (e3 = op(op(e2, e3), e2)) & ~ (e3 = op(e2, op(e2, e3)))) | ~ sP11), inference(usedef, [], [e65])).
fof(e65, plain, (sP11 <=> (~ (e2 = op(op(e2, e3), e3)) & (e3 = op(op(e2, e3), e2)) & ~ (e3 = op(e2, op(e2, e3))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f66, plain, ((~ (e3 = op(op(e3, e0), e0)) & (e0 = op(op(e3, e0), e3)) & ~ (e0 = op(e3, op(e3, e0)))) | ~ sP12), inference(usedef, [], [e66])).
fof(e66, plain, (sP12 <=> (~ (e3 = op(op(e3, e0), e0)) & (e0 = op(op(e3, e0), e3)) & ~ (e0 = op(e3, op(e3, e0))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f67, plain, ((~ (e3 = op(op(e3, e1), e1)) & (e1 = op(op(e3, e1), e3)) & ~ (e1 = op(e3, op(e3, e1)))) | ~ sP13), inference(usedef, [], [e67])).
fof(e67, plain, (sP13 <=> (~ (e3 = op(op(e3, e1), e1)) & (e1 = op(op(e3, e1), e3)) & ~ (e1 = op(e3, op(e3, e1))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f68, plain, ((~ (e3 = op(op(e3, e2), e2)) & (e2 = op(op(e3, e2), e3)) & ~ (e2 = op(e3, op(e3, e2)))) | ~ sP14), inference(usedef, [], [e68])).
fof(e68, plain, (sP14 <=> (~ (e3 = op(op(e3, e2), e2)) & (e2 = op(op(e3, e2), e3)) & ~ (e2 = op(e3, op(e3, e2))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f69, plain, (((e3 = op(e0, op(e0, e3))) & (e2 = op(e0, op(e0, e2))) & (e1 = op(e0, op(e0, e1))) & (e0 = op(e0, op(e0, e0))) & (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP15), inference(usedef, [], [e69])).
fof(e69, plain, (sP15 <=> ((e3 = op(e0, op(e0, e3))) & (e2 = op(e0, op(e0, e2))) & (e1 = op(e0, op(e0, e1))) & (e0 = op(e0, op(e0, e0))) & (e0 = op(e0, e0)) & (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f70, plain, (((e3 = op(e0, op(e0, e3))) & (e2 = op(e0, op(e0, e2))) & (e1 = op(e0, op(e0, e1))) & (e0 = op(e0, op(e0, e0))) & (e0 = op(e1, e0)) & (e0 = op(e0, e1))) | ~ sP16), inference(usedef, [], [e70])).
fof(e70, plain, (sP16 <=> ((e3 = op(e0, op(e0, e3))) & (e2 = op(e0, op(e0, e2))) & (e1 = op(e0, op(e0, e1))) & (e0 = op(e0, op(e0, e0))) & (e0 = op(e1, e0)) & (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f71, plain, (((e3 = op(e0, op(e0, e3))) & (e2 = op(e0, op(e0, e2))) & (e1 = op(e0, op(e0, e1))) & (e0 = op(e0, op(e0, e0))) & (e0 = op(e2, e0)) & (e0 = op(e0, e2))) | ~ sP17), inference(usedef, [], [e71])).
fof(e71, plain, (sP17 <=> ((e3 = op(e0, op(e0, e3))) & (e2 = op(e0, op(e0, e2))) & (e1 = op(e0, op(e0, e1))) & (e0 = op(e0, op(e0, e0))) & (e0 = op(e2, e0)) & (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f72, plain, (((e3 = op(e0, op(e0, e3))) & (e2 = op(e0, op(e0, e2))) & (e1 = op(e0, op(e0, e1))) & (e0 = op(e0, op(e0, e0))) & (e0 = op(e3, e0)) & (e0 = op(e0, e3))) | ~ sP18), inference(usedef, [], [e72])).
fof(e72, plain, (sP18 <=> ((e3 = op(e0, op(e0, e3))) & (e2 = op(e0, op(e0, e2))) & (e1 = op(e0, op(e0, e1))) & (e0 = op(e0, op(e0, e0))) & (e0 = op(e3, e0)) & (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f73, plain, (((e3 = op(e1, op(e1, e3))) & (e2 = op(e1, op(e1, e2))) & (e1 = op(e1, op(e1, e1))) & (e0 = op(e1, op(e1, e0))) & (e1 = op(e0, e1)) & (e1 = op(e1, e0))) | ~ sP19), inference(usedef, [], [e73])).
fof(e73, plain, (sP19 <=> ((e3 = op(e1, op(e1, e3))) & (e2 = op(e1, op(e1, e2))) & (e1 = op(e1, op(e1, e1))) & (e0 = op(e1, op(e1, e0))) & (e1 = op(e0, e1)) & (e1 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP19])])).
fof(f74, plain, (((e3 = op(e1, op(e1, e3))) & (e2 = op(e1, op(e1, e2))) & (e1 = op(e1, op(e1, e1))) & (e0 = op(e1, op(e1, e0))) & (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP20), inference(usedef, [], [e74])).
fof(e74, plain, (sP20 <=> ((e3 = op(e1, op(e1, e3))) & (e2 = op(e1, op(e1, e2))) & (e1 = op(e1, op(e1, e1))) & (e0 = op(e1, op(e1, e0))) & (e1 = op(e1, e1)) & (e1 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP20])])).
fof(f75, plain, (((e3 = op(e1, op(e1, e3))) & (e2 = op(e1, op(e1, e2))) & (e1 = op(e1, op(e1, e1))) & (e0 = op(e1, op(e1, e0))) & (e1 = op(e2, e1)) & (e1 = op(e1, e2))) | ~ sP21), inference(usedef, [], [e75])).
fof(e75, plain, (sP21 <=> ((e3 = op(e1, op(e1, e3))) & (e2 = op(e1, op(e1, e2))) & (e1 = op(e1, op(e1, e1))) & (e0 = op(e1, op(e1, e0))) & (e1 = op(e2, e1)) & (e1 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP21])])).
fof(f76, plain, (((e3 = op(e1, op(e1, e3))) & (e2 = op(e1, op(e1, e2))) & (e1 = op(e1, op(e1, e1))) & (e0 = op(e1, op(e1, e0))) & (e1 = op(e3, e1)) & (e1 = op(e1, e3))) | ~ sP22), inference(usedef, [], [e76])).
fof(e76, plain, (sP22 <=> ((e3 = op(e1, op(e1, e3))) & (e2 = op(e1, op(e1, e2))) & (e1 = op(e1, op(e1, e1))) & (e0 = op(e1, op(e1, e0))) & (e1 = op(e3, e1)) & (e1 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP22])])).
fof(f77, plain, (((e3 = op(e2, op(e2, e3))) & (e2 = op(e2, op(e2, e2))) & (e1 = op(e2, op(e2, e1))) & (e0 = op(e2, op(e2, e0))) & (e2 = op(e0, e2)) & (e2 = op(e2, e0))) | ~ sP23), inference(usedef, [], [e77])).
fof(e77, plain, (sP23 <=> ((e3 = op(e2, op(e2, e3))) & (e2 = op(e2, op(e2, e2))) & (e1 = op(e2, op(e2, e1))) & (e0 = op(e2, op(e2, e0))) & (e2 = op(e0, e2)) & (e2 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP23])])).
fof(f78, plain, (((e3 = op(e2, op(e2, e3))) & (e2 = op(e2, op(e2, e2))) & (e1 = op(e2, op(e2, e1))) & (e0 = op(e2, op(e2, e0))) & (e2 = op(e1, e2)) & (e2 = op(e2, e1))) | ~ sP24), inference(usedef, [], [e78])).
fof(e78, plain, (sP24 <=> ((e3 = op(e2, op(e2, e3))) & (e2 = op(e2, op(e2, e2))) & (e1 = op(e2, op(e2, e1))) & (e0 = op(e2, op(e2, e0))) & (e2 = op(e1, e2)) & (e2 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP24])])).
fof(f79, plain, (((e3 = op(e2, op(e2, e3))) & (e2 = op(e2, op(e2, e2))) & (e1 = op(e2, op(e2, e1))) & (e0 = op(e2, op(e2, e0))) & (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP25), inference(usedef, [], [e79])).
fof(e79, plain, (sP25 <=> ((e3 = op(e2, op(e2, e3))) & (e2 = op(e2, op(e2, e2))) & (e1 = op(e2, op(e2, e1))) & (e0 = op(e2, op(e2, e0))) & (e2 = op(e2, e2)) & (e2 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP25])])).
fof(f80, plain, (((e3 = op(e2, op(e2, e3))) & (e2 = op(e2, op(e2, e2))) & (e1 = op(e2, op(e2, e1))) & (e0 = op(e2, op(e2, e0))) & (e2 = op(e3, e2)) & (e2 = op(e2, e3))) | ~ sP26), inference(usedef, [], [e80])).
fof(e80, plain, (sP26 <=> ((e3 = op(e2, op(e2, e3))) & (e2 = op(e2, op(e2, e2))) & (e1 = op(e2, op(e2, e1))) & (e0 = op(e2, op(e2, e0))) & (e2 = op(e3, e2)) & (e2 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP26])])).
fof(f81, plain, (((e3 = op(e3, op(e3, e3))) & (e2 = op(e3, op(e3, e2))) & (e1 = op(e3, op(e3, e1))) & (e0 = op(e3, op(e3, e0))) & (e3 = op(e0, e3)) & (e3 = op(e3, e0))) | ~ sP27), inference(usedef, [], [e81])).
fof(e81, plain, (sP27 <=> ((e3 = op(e3, op(e3, e3))) & (e2 = op(e3, op(e3, e2))) & (e1 = op(e3, op(e3, e1))) & (e0 = op(e3, op(e3, e0))) & (e3 = op(e0, e3)) & (e3 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP27])])).
fof(f82, plain, (((e3 = op(e3, op(e3, e3))) & (e2 = op(e3, op(e3, e2))) & (e1 = op(e3, op(e3, e1))) & (e0 = op(e3, op(e3, e0))) & (e3 = op(e1, e3)) & (e3 = op(e3, e1))) | ~ sP28), inference(usedef, [], [e82])).
fof(e82, plain, (sP28 <=> ((e3 = op(e3, op(e3, e3))) & (e2 = op(e3, op(e3, e2))) & (e1 = op(e3, op(e3, e1))) & (e0 = op(e3, op(e3, e0))) & (e3 = op(e1, e3)) & (e3 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP28])])).
fof(f83, plain, (((e3 = op(e3, op(e3, e3))) & (e2 = op(e3, op(e3, e2))) & (e1 = op(e3, op(e3, e1))) & (e0 = op(e3, op(e3, e0))) & (e3 = op(e2, e3)) & (e3 = op(e3, e2))) | ~ sP29), inference(usedef, [], [e83])).
fof(e83, plain, (sP29 <=> ((e3 = op(e3, op(e3, e3))) & (e2 = op(e3, op(e3, e2))) & (e1 = op(e3, op(e3, e1))) & (e0 = op(e3, op(e3, e0))) & (e3 = op(e2, e3)) & (e3 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP29])])).
fof(f5, plain, ((((e3 = op(e3, op(e3, e3))) & (e2 = op(e3, op(e3, e2))) & (e1 = op(e3, op(e3, e1))) & (e0 = op(e3, op(e3, e0))) & (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ((e3 = op(e3, op(e3, e3))) & (e2 = op(e3, op(e3, e2))) & (e1 = op(e3, op(e3, e1))) & (e0 = op(e3, op(e3, e0))) & (e3 = op(e2, e3)) & (e3 = op(e3, e2))) | ((e3 = op(e3, op(e3, e3))) & (e2 = op(e3, op(e3, e2))) & (e1 = op(e3, op(e3, e1))) & (e0 = op(e3, op(e3, e0))) & (e3 = op(e1, e3)) & (e3 = op(e3, e1))) | ((e3 = op(e3, op(e3, e3))) & (e2 = op(e3, op(e3, e2))) & (e1 = op(e3, op(e3, e1))) & (e0 = op(e3, op(e3, e0))) & (e3 = op(e0, e3)) & (e3 = op(e3, e0))) | ((e3 = op(e2, op(e2, e3))) & (e2 = op(e2, op(e2, e2))) & (e1 = op(e2, op(e2, e1))) & (e0 = op(e2, op(e2, e0))) & (e2 = op(e3, e2)) & (e2 = op(e2, e3))) | ((e3 = op(e2, op(e2, e3))) & (e2 = op(e2, op(e2, e2))) & (e1 = op(e2, op(e2, e1))) & (e0 = op(e2, op(e2, e0))) & (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ((e3 = op(e2, op(e2, e3))) & (e2 = op(e2, op(e2, e2))) & (e1 = op(e2, op(e2, e1))) & (e0 = op(e2, op(e2, e0))) & (e2 = op(e1, e2)) & (e2 = op(e2, e1))) | ((e3 = op(e2, op(e2, e3))) & (e2 = op(e2, op(e2, e2))) & (e1 = op(e2, op(e2, e1))) & (e0 = op(e2, op(e2, e0))) & (e2 = op(e0, e2)) & (e2 = op(e2, e0))) | ((e3 = op(e1, op(e1, e3))) & (e2 = op(e1, op(e1, e2))) & (e1 = op(e1, op(e1, e1))) & (e0 = op(e1, op(e1, e0))) & (e1 = op(e3, e1)) & (e1 = op(e1, e3))) | ((e3 = op(e1, op(e1, e3))) & (e2 = op(e1, op(e1, e2))) & (e1 = op(e1, op(e1, e1))) & (e0 = op(e1, op(e1, e0))) & (e1 = op(e2, e1)) & (e1 = op(e1, e2))) | ((e3 = op(e1, op(e1, e3))) & (e2 = op(e1, op(e1, e2))) & (e1 = op(e1, op(e1, e1))) & (e0 = op(e1, op(e1, e0))) & (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ((e3 = op(e1, op(e1, e3))) & (e2 = op(e1, op(e1, e2))) & (e1 = op(e1, op(e1, e1))) & (e0 = op(e1, op(e1, e0))) & (e1 = op(e0, e1)) & (e1 = op(e1, e0))) | ((e3 = op(e0, op(e0, e3))) & (e2 = op(e0, op(e0, e2))) & (e1 = op(e0, op(e0, e1))) & (e0 = op(e0, op(e0, e0))) & (e0 = op(e3, e0)) & (e0 = op(e0, e3))) | ((e3 = op(e0, op(e0, e3))) & (e2 = op(e0, op(e0, e2))) & (e1 = op(e0, op(e0, e1))) & (e0 = op(e0, op(e0, e0))) & (e0 = op(e2, e0)) & (e0 = op(e0, e2))) | ((e3 = op(e0, op(e0, e3))) & (e2 = op(e0, op(e0, e2))) & (e1 = op(e0, op(e0, e1))) & (e0 = op(e0, op(e0, e0))) & (e0 = op(e1, e0)) & (e0 = op(e0, e1))) | ((e3 = op(e0, op(e0, e3))) & (e2 = op(e0, op(e0, e2))) & (e1 = op(e0, op(e0, e1))) & (e0 = op(e0, op(e0, e0))) & (e0 = op(e0, e0)) & (e0 = op(e0, e0)))) & ((~ (e3 = op(op(e3, e3), e3)) & (e3 = op(op(e3, e3), e3)) & ~ (e3 = op(e3, op(e3, e3)))) | (~ (e3 = op(op(e3, e2), e2)) & (e2 = op(op(e3, e2), e3)) & ~ (e2 = op(e3, op(e3, e2)))) | (~ (e3 = op(op(e3, e1), e1)) & (e1 = op(op(e3, e1), e3)) & ~ (e1 = op(e3, op(e3, e1)))) | (~ (e3 = op(op(e3, e0), e0)) & (e0 = op(op(e3, e0), e3)) & ~ (e0 = op(e3, op(e3, e0)))) | (~ (e2 = op(op(e2, e3), e3)) & (e3 = op(op(e2, e3), e2)) & ~ (e3 = op(e2, op(e2, e3)))) | (~ (e2 = op(op(e2, e2), e2)) & (e2 = op(op(e2, e2), e2)) & ~ (e2 = op(e2, op(e2, e2)))) | (~ (e2 = op(op(e2, e1), e1)) & (e1 = op(op(e2, e1), e2)) & ~ (e1 = op(e2, op(e2, e1)))) | (~ (e2 = op(op(e2, e0), e0)) & (e0 = op(op(e2, e0), e2)) & ~ (e0 = op(e2, op(e2, e0)))) | (~ (e1 = op(op(e1, e3), e3)) & (e3 = op(op(e1, e3), e1)) & ~ (e3 = op(e1, op(e1, e3)))) | (~ (e1 = op(op(e1, e2), e2)) & (e2 = op(op(e1, e2), e1)) & ~ (e2 = op(e1, op(e1, e2)))) | (~ (e1 = op(op(e1, e1), e1)) & (e1 = op(op(e1, e1), e1)) & ~ (e1 = op(e1, op(e1, e1)))) | (~ (e1 = op(op(e1, e0), e0)) & (e0 = op(op(e1, e0), e1)) & ~ (e0 = op(e1, op(e1, e0)))) | (~ (e0 = op(op(e0, e3), e3)) & (e3 = op(op(e0, e3), e0)) & ~ (e3 = op(e0, op(e0, e3)))) | (~ (e0 = op(op(e0, e2), e2)) & (e2 = op(op(e0, e2), e0)) & ~ (e2 = op(e0, op(e0, e2)))) | (~ (e0 = op(op(e0, e1), e1)) & (e1 = op(op(e0, e1), e0)) & ~ (e1 = op(e0, op(e0, e1)))) | (~ (e0 = op(op(e0, e0), e0)) & (e0 = op(op(e0, e0), e0)) & ~ (e0 = op(e0, op(e0, e0)))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG164+1.p', ax5)).
fof(f1127, plain, (spl30_136 | spl30_133 | spl30_130 | spl30_127 | spl30_124 | spl30_122 | spl30_119 | spl30_116 | spl30_113 | spl30_110 | spl30_108 | spl30_105 | spl30_102 | spl30_99 | spl30_96 | spl30_138), inference(avatar_split_clause, [], [f353, f1123, f904, f919, f934, f949, f964, f975, f990, f1005, f1020, f1035, f1046, f1061, f1076, f1091, f1106])).
fof(f1123, plain, (spl30_138 <=> (e3 = op(op(e3, e3), e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_138])])).
fof(f353, plain, ((e3 = op(op(e3, e3), e3)) | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f84])).
fof(f1126, plain, (spl30_136 | spl30_133 | spl30_130 | spl30_127 | spl30_124 | spl30_122 | spl30_119 | spl30_116 | spl30_113 | spl30_110 | spl30_108 | spl30_105 | spl30_102 | spl30_99 | spl30_96 | ~ spl30_138), inference(avatar_split_clause, [], [f354, f1123, f904, f919, f934, f949, f964, f975, f990, f1005, f1020, f1035, f1046, f1061, f1076, f1091, f1106])).
fof(f354, plain, (~ (e3 = op(op(e3, e3), e3)) | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f84])).
fof(f1121, plain, (spl30_95 | spl30_94 | spl30_93 | spl30_88 | spl30_87 | spl30_86 | spl30_85 | spl30_80 | spl30_79 | spl30_78 | spl30_77 | spl30_72 | spl30_71 | spl30_70 | spl30_65 | spl30_4), inference(avatar_split_clause, [], [f355, f398, f690, f716, f726, f736, f762, f772, f782, f792, f818, f828, f838, f848, f874, f884, f894])).
fof(f894, plain, (spl30_95 <=> sP15), introduced(avatar_definition, [new_symbols(naming, [spl30_95])])).
fof(f884, plain, (spl30_94 <=> sP16), introduced(avatar_definition, [new_symbols(naming, [spl30_94])])).
fof(f874, plain, (spl30_93 <=> sP17), introduced(avatar_definition, [new_symbols(naming, [spl30_93])])).
fof(f848, plain, (spl30_88 <=> sP18), introduced(avatar_definition, [new_symbols(naming, [spl30_88])])).
fof(f838, plain, (spl30_87 <=> sP19), introduced(avatar_definition, [new_symbols(naming, [spl30_87])])).
fof(f828, plain, (spl30_86 <=> sP20), introduced(avatar_definition, [new_symbols(naming, [spl30_86])])).
fof(f818, plain, (spl30_85 <=> sP21), introduced(avatar_definition, [new_symbols(naming, [spl30_85])])).
fof(f792, plain, (spl30_80 <=> sP22), introduced(avatar_definition, [new_symbols(naming, [spl30_80])])).
fof(f782, plain, (spl30_79 <=> sP23), introduced(avatar_definition, [new_symbols(naming, [spl30_79])])).
fof(f772, plain, (spl30_78 <=> sP24), introduced(avatar_definition, [new_symbols(naming, [spl30_78])])).
fof(f762, plain, (spl30_77 <=> sP25), introduced(avatar_definition, [new_symbols(naming, [spl30_77])])).
fof(f736, plain, (spl30_72 <=> sP26), introduced(avatar_definition, [new_symbols(naming, [spl30_72])])).
fof(f726, plain, (spl30_71 <=> sP27), introduced(avatar_definition, [new_symbols(naming, [spl30_71])])).
fof(f716, plain, (spl30_70 <=> sP28), introduced(avatar_definition, [new_symbols(naming, [spl30_70])])).
fof(f690, plain, (spl30_65 <=> sP29), introduced(avatar_definition, [new_symbols(naming, [spl30_65])])).
fof(f355, plain, ((e3 = op(e3, e3)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15), inference(cnf_transformation, [], [f84])).
fof(f1119, plain, (spl30_95 | spl30_94 | spl30_93 | spl30_88 | spl30_87 | spl30_86 | spl30_85 | spl30_80 | spl30_79 | spl30_78 | spl30_77 | spl30_72 | spl30_71 | spl30_70 | spl30_65 | spl30_69), inference(avatar_split_clause, [], [f357, f709, f690, f716, f726, f736, f762, f772, f782, f792, f818, f828, f838, f848, f874, f884, f894])).
fof(f357, plain, ((e0 = op(e3, op(e3, e0))) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15), inference(cnf_transformation, [], [f84])).
fof(f1118, plain, (spl30_95 | spl30_94 | spl30_93 | spl30_88 | spl30_87 | spl30_86 | spl30_85 | spl30_80 | spl30_79 | spl30_78 | spl30_77 | spl30_72 | spl30_71 | spl30_70 | spl30_65 | spl30_68), inference(avatar_split_clause, [], [f358, f704, f690, f716, f726, f736, f762, f772, f782, f792, f818, f828, f838, f848, f874, f884, f894])).
fof(f358, plain, ((e1 = op(e3, op(e3, e1))) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15), inference(cnf_transformation, [], [f84])).
fof(f1117, plain, (spl30_95 | spl30_94 | spl30_93 | spl30_88 | spl30_87 | spl30_86 | spl30_85 | spl30_80 | spl30_79 | spl30_78 | spl30_77 | spl30_72 | spl30_71 | spl30_70 | spl30_65 | spl30_67), inference(avatar_split_clause, [], [f359, f699, f690, f716, f726, f736, f762, f772, f782, f792, f818, f828, f838, f848, f874, f884, f894])).
fof(f359, plain, ((e2 = op(e3, op(e3, e2))) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15), inference(cnf_transformation, [], [f84])).
fof(f1116, plain, (spl30_95 | spl30_94 | spl30_93 | spl30_88 | spl30_87 | spl30_86 | spl30_85 | spl30_80 | spl30_79 | spl30_78 | spl30_77 | spl30_72 | spl30_71 | spl30_70 | spl30_65 | spl30_66), inference(avatar_split_clause, [], [f360, f694, f690, f716, f726, f736, f762, f772, f782, f792, f818, f828, f838, f848, f874, f884, f894])).
fof(f360, plain, ((e3 = op(e3, op(e3, e3))) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15), inference(cnf_transformation, [], [f84])).
fof(f1115, plain, (~ spl30_136 | ~ spl30_92), inference(avatar_split_clause, [], [f349, f867, f1106])).
fof(f349, plain, (~ (e0 = op(e0, op(e0, e0))) | ~ sP0), inference(cnf_transformation, [], [f114])).
fof(f114, plain, ((~ (e0 = op(op(e0, e0), e0)) & (e0 = op(op(e0, e0), e0)) & ~ (e0 = op(e0, op(e0, e0)))) | ~ sP0), inference(nnf_transformation, [], [f54])).
fof(f1114, plain, (~ spl30_136 | spl30_137), inference(avatar_split_clause, [], [f350, f1110, f1106])).
fof(f1110, plain, (spl30_137 <=> (e0 = op(op(e0, e0), e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_137])])).
fof(f350, plain, ((e0 = op(op(e0, e0), e0)) | ~ sP0), inference(cnf_transformation, [], [f114])).
fof(f1113, plain, (~ spl30_136 | ~ spl30_137), inference(avatar_split_clause, [], [f351, f1110, f1106])).
fof(f351, plain, (~ (e0 = op(op(e0, e0), e0)) | ~ sP0), inference(cnf_transformation, [], [f114])).
fof(f1104, plain, (~ spl30_133 | ~ spl30_91), inference(avatar_split_clause, [], [f346, f862, f1091])).
fof(f346, plain, (~ (e1 = op(e0, op(e0, e1))) | ~ sP1), inference(cnf_transformation, [], [f113])).
fof(f113, plain, ((~ (e0 = op(op(e0, e1), e1)) & (e1 = op(op(e0, e1), e0)) & ~ (e1 = op(e0, op(e0, e1)))) | ~ sP1), inference(nnf_transformation, [], [f55])).
fof(f1103, plain, (~ spl30_133 | spl30_135), inference(avatar_split_clause, [], [f347, f1100, f1091])).
fof(f347, plain, ((e1 = op(op(e0, e1), e0)) | ~ sP1), inference(cnf_transformation, [], [f113])).
fof(f1098, plain, (~ spl30_133 | ~ spl30_134), inference(avatar_split_clause, [], [f348, f1095, f1091])).
fof(f348, plain, (~ (e0 = op(op(e0, e1), e1)) | ~ sP1), inference(cnf_transformation, [], [f113])).
fof(f1089, plain, (~ spl30_130 | ~ spl30_90), inference(avatar_split_clause, [], [f343, f857, f1076])).
fof(f343, plain, (~ (e2 = op(e0, op(e0, e2))) | ~ sP2), inference(cnf_transformation, [], [f112])).
fof(f112, plain, ((~ (e0 = op(op(e0, e2), e2)) & (e2 = op(op(e0, e2), e0)) & ~ (e2 = op(e0, op(e0, e2)))) | ~ sP2), inference(nnf_transformation, [], [f56])).
fof(f1088, plain, (~ spl30_130 | spl30_132), inference(avatar_split_clause, [], [f344, f1085, f1076])).
fof(f344, plain, ((e2 = op(op(e0, e2), e0)) | ~ sP2), inference(cnf_transformation, [], [f112])).
fof(f1083, plain, (~ spl30_130 | ~ spl30_131), inference(avatar_split_clause, [], [f345, f1080, f1076])).
fof(f345, plain, (~ (e0 = op(op(e0, e2), e2)) | ~ sP2), inference(cnf_transformation, [], [f112])).
fof(f1074, plain, (~ spl30_127 | ~ spl30_89), inference(avatar_split_clause, [], [f340, f852, f1061])).
fof(f340, plain, (~ (e3 = op(e0, op(e0, e3))) | ~ sP3), inference(cnf_transformation, [], [f111])).
fof(f111, plain, ((~ (e0 = op(op(e0, e3), e3)) & (e3 = op(op(e0, e3), e0)) & ~ (e3 = op(e0, op(e0, e3)))) | ~ sP3), inference(nnf_transformation, [], [f57])).
fof(f1073, plain, (~ spl30_127 | spl30_129), inference(avatar_split_clause, [], [f341, f1070, f1061])).
fof(f341, plain, ((e3 = op(op(e0, e3), e0)) | ~ sP3), inference(cnf_transformation, [], [f111])).
fof(f1068, plain, (~ spl30_127 | ~ spl30_128), inference(avatar_split_clause, [], [f342, f1065, f1061])).
fof(f342, plain, (~ (e0 = op(op(e0, e3), e3)) | ~ sP3), inference(cnf_transformation, [], [f111])).
fof(f1059, plain, (~ spl30_124 | ~ spl30_84), inference(avatar_split_clause, [], [f337, f811, f1046])).
fof(f337, plain, (~ (e0 = op(e1, op(e1, e0))) | ~ sP4), inference(cnf_transformation, [], [f110])).
fof(f110, plain, ((~ (e1 = op(op(e1, e0), e0)) & (e0 = op(op(e1, e0), e1)) & ~ (e0 = op(e1, op(e1, e0)))) | ~ sP4), inference(nnf_transformation, [], [f58])).
fof(f1058, plain, (~ spl30_124 | spl30_126), inference(avatar_split_clause, [], [f338, f1055, f1046])).
fof(f338, plain, ((e0 = op(op(e1, e0), e1)) | ~ sP4), inference(cnf_transformation, [], [f110])).
fof(f1053, plain, (~ spl30_124 | ~ spl30_125), inference(avatar_split_clause, [], [f339, f1050, f1046])).
fof(f339, plain, (~ (e1 = op(op(e1, e0), e0)) | ~ sP4), inference(cnf_transformation, [], [f110])).
fof(f1043, plain, (~ spl30_122 | spl30_123), inference(avatar_split_clause, [], [f335, f1039, f1035])).
fof(f1039, plain, (spl30_123 <=> (e1 = op(op(e1, e1), e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_123])])).
fof(f335, plain, ((e1 = op(op(e1, e1), e1)) | ~ sP5), inference(cnf_transformation, [], [f109])).
fof(f109, plain, ((~ (e1 = op(op(e1, e1), e1)) & (e1 = op(op(e1, e1), e1)) & ~ (e1 = op(e1, op(e1, e1)))) | ~ sP5), inference(nnf_transformation, [], [f59])).
fof(f1042, plain, (~ spl30_122 | ~ spl30_123), inference(avatar_split_clause, [], [f336, f1039, f1035])).
fof(f336, plain, (~ (e1 = op(op(e1, e1), e1)) | ~ sP5), inference(cnf_transformation, [], [f109])).
fof(f1033, plain, (~ spl30_119 | ~ spl30_82), inference(avatar_split_clause, [], [f331, f801, f1020])).
fof(f331, plain, (~ (e2 = op(e1, op(e1, e2))) | ~ sP6), inference(cnf_transformation, [], [f108])).
fof(f108, plain, ((~ (e1 = op(op(e1, e2), e2)) & (e2 = op(op(e1, e2), e1)) & ~ (e2 = op(e1, op(e1, e2)))) | ~ sP6), inference(nnf_transformation, [], [f60])).
fof(f1032, plain, (~ spl30_119 | spl30_121), inference(avatar_split_clause, [], [f332, f1029, f1020])).
fof(f332, plain, ((e2 = op(op(e1, e2), e1)) | ~ sP6), inference(cnf_transformation, [], [f108])).
fof(f1027, plain, (~ spl30_119 | ~ spl30_120), inference(avatar_split_clause, [], [f333, f1024, f1020])).
fof(f333, plain, (~ (e1 = op(op(e1, e2), e2)) | ~ sP6), inference(cnf_transformation, [], [f108])).
fof(f1018, plain, (~ spl30_116 | ~ spl30_81), inference(avatar_split_clause, [], [f328, f796, f1005])).
fof(f328, plain, (~ (e3 = op(e1, op(e1, e3))) | ~ sP7), inference(cnf_transformation, [], [f107])).
fof(f107, plain, ((~ (e1 = op(op(e1, e3), e3)) & (e3 = op(op(e1, e3), e1)) & ~ (e3 = op(e1, op(e1, e3)))) | ~ sP7), inference(nnf_transformation, [], [f61])).
fof(f1017, plain, (~ spl30_116 | spl30_118), inference(avatar_split_clause, [], [f329, f1014, f1005])).
fof(f329, plain, ((e3 = op(op(e1, e3), e1)) | ~ sP7), inference(cnf_transformation, [], [f107])).
fof(f1012, plain, (~ spl30_116 | ~ spl30_117), inference(avatar_split_clause, [], [f330, f1009, f1005])).
fof(f330, plain, (~ (e1 = op(op(e1, e3), e3)) | ~ sP7), inference(cnf_transformation, [], [f107])).
fof(f1003, plain, (~ spl30_113 | ~ spl30_76), inference(avatar_split_clause, [], [f325, f755, f990])).
fof(f325, plain, (~ (e0 = op(e2, op(e2, e0))) | ~ sP8), inference(cnf_transformation, [], [f106])).
fof(f106, plain, ((~ (e2 = op(op(e2, e0), e0)) & (e0 = op(op(e2, e0), e2)) & ~ (e0 = op(e2, op(e2, e0)))) | ~ sP8), inference(nnf_transformation, [], [f62])).
fof(f1002, plain, (~ spl30_113 | spl30_115), inference(avatar_split_clause, [], [f326, f999, f990])).
fof(f326, plain, ((e0 = op(op(e2, e0), e2)) | ~ sP8), inference(cnf_transformation, [], [f106])).
fof(f997, plain, (~ spl30_113 | ~ spl30_114), inference(avatar_split_clause, [], [f327, f994, f990])).
fof(f327, plain, (~ (e2 = op(op(e2, e0), e0)) | ~ sP8), inference(cnf_transformation, [], [f106])).
fof(f988, plain, (~ spl30_110 | ~ spl30_75), inference(avatar_split_clause, [], [f322, f750, f975])).
fof(f322, plain, (~ (e1 = op(e2, op(e2, e1))) | ~ sP9), inference(cnf_transformation, [], [f105])).
fof(f105, plain, ((~ (e2 = op(op(e2, e1), e1)) & (e1 = op(op(e2, e1), e2)) & ~ (e1 = op(e2, op(e2, e1)))) | ~ sP9), inference(nnf_transformation, [], [f63])).
fof(f987, plain, (~ spl30_110 | spl30_112), inference(avatar_split_clause, [], [f323, f984, f975])).
fof(f323, plain, ((e1 = op(op(e2, e1), e2)) | ~ sP9), inference(cnf_transformation, [], [f105])).
fof(f982, plain, (~ spl30_110 | ~ spl30_111), inference(avatar_split_clause, [], [f324, f979, f975])).
fof(f324, plain, (~ (e2 = op(op(e2, e1), e1)) | ~ sP9), inference(cnf_transformation, [], [f105])).
fof(f973, plain, (~ spl30_108 | ~ spl30_74), inference(avatar_split_clause, [], [f319, f745, f964])).
fof(f319, plain, (~ (e2 = op(e2, op(e2, e2))) | ~ sP10), inference(cnf_transformation, [], [f104])).
fof(f104, plain, ((~ (e2 = op(op(e2, e2), e2)) & (e2 = op(op(e2, e2), e2)) & ~ (e2 = op(e2, op(e2, e2)))) | ~ sP10), inference(nnf_transformation, [], [f64])).
fof(f972, plain, (~ spl30_108 | spl30_109), inference(avatar_split_clause, [], [f320, f968, f964])).
fof(f968, plain, (spl30_109 <=> (e2 = op(op(e2, e2), e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_109])])).
fof(f320, plain, ((e2 = op(op(e2, e2), e2)) | ~ sP10), inference(cnf_transformation, [], [f104])).
fof(f971, plain, (~ spl30_108 | ~ spl30_109), inference(avatar_split_clause, [], [f321, f968, f964])).
fof(f321, plain, (~ (e2 = op(op(e2, e2), e2)) | ~ sP10), inference(cnf_transformation, [], [f104])).
fof(f962, plain, (~ spl30_105 | ~ spl30_73), inference(avatar_split_clause, [], [f316, f740, f949])).
fof(f316, plain, (~ (e3 = op(e2, op(e2, e3))) | ~ sP11), inference(cnf_transformation, [], [f103])).
fof(f103, plain, ((~ (e2 = op(op(e2, e3), e3)) & (e3 = op(op(e2, e3), e2)) & ~ (e3 = op(e2, op(e2, e3)))) | ~ sP11), inference(nnf_transformation, [], [f65])).
fof(f961, plain, (~ spl30_105 | spl30_107), inference(avatar_split_clause, [], [f317, f958, f949])).
fof(f317, plain, ((e3 = op(op(e2, e3), e2)) | ~ sP11), inference(cnf_transformation, [], [f103])).
fof(f956, plain, (~ spl30_105 | ~ spl30_106), inference(avatar_split_clause, [], [f318, f953, f949])).
fof(f318, plain, (~ (e2 = op(op(e2, e3), e3)) | ~ sP11), inference(cnf_transformation, [], [f103])).
fof(f947, plain, (~ spl30_102 | ~ spl30_69), inference(avatar_split_clause, [], [f313, f709, f934])).
fof(f313, plain, (~ (e0 = op(e3, op(e3, e0))) | ~ sP12), inference(cnf_transformation, [], [f102])).
fof(f102, plain, ((~ (e3 = op(op(e3, e0), e0)) & (e0 = op(op(e3, e0), e3)) & ~ (e0 = op(e3, op(e3, e0)))) | ~ sP12), inference(nnf_transformation, [], [f66])).
fof(f946, plain, (~ spl30_102 | spl30_104), inference(avatar_split_clause, [], [f314, f943, f934])).
fof(f314, plain, ((e0 = op(op(e3, e0), e3)) | ~ sP12), inference(cnf_transformation, [], [f102])).
fof(f941, plain, (~ spl30_102 | ~ spl30_103), inference(avatar_split_clause, [], [f315, f938, f934])).
fof(f315, plain, (~ (e3 = op(op(e3, e0), e0)) | ~ sP12), inference(cnf_transformation, [], [f102])).
fof(f932, plain, (~ spl30_99 | ~ spl30_68), inference(avatar_split_clause, [], [f310, f704, f919])).
fof(f310, plain, (~ (e1 = op(e3, op(e3, e1))) | ~ sP13), inference(cnf_transformation, [], [f101])).
fof(f101, plain, ((~ (e3 = op(op(e3, e1), e1)) & (e1 = op(op(e3, e1), e3)) & ~ (e1 = op(e3, op(e3, e1)))) | ~ sP13), inference(nnf_transformation, [], [f67])).
fof(f931, plain, (~ spl30_99 | spl30_101), inference(avatar_split_clause, [], [f311, f928, f919])).
fof(f311, plain, ((e1 = op(op(e3, e1), e3)) | ~ sP13), inference(cnf_transformation, [], [f101])).
fof(f926, plain, (~ spl30_99 | ~ spl30_100), inference(avatar_split_clause, [], [f312, f923, f919])).
fof(f312, plain, (~ (e3 = op(op(e3, e1), e1)) | ~ sP13), inference(cnf_transformation, [], [f101])).
fof(f917, plain, (~ spl30_96 | ~ spl30_67), inference(avatar_split_clause, [], [f307, f699, f904])).
fof(f307, plain, (~ (e2 = op(e3, op(e3, e2))) | ~ sP14), inference(cnf_transformation, [], [f100])).
fof(f100, plain, ((~ (e3 = op(op(e3, e2), e2)) & (e2 = op(op(e3, e2), e3)) & ~ (e2 = op(e3, op(e3, e2)))) | ~ sP14), inference(nnf_transformation, [], [f68])).
fof(f916, plain, (~ spl30_96 | spl30_98), inference(avatar_split_clause, [], [f308, f913, f904])).
fof(f308, plain, ((e2 = op(op(e3, e2), e3)) | ~ sP14), inference(cnf_transformation, [], [f100])).
fof(f911, plain, (~ spl30_96 | ~ spl30_97), inference(avatar_split_clause, [], [f309, f908, f904])).
fof(f309, plain, (~ (e3 = op(op(e3, e2), e2)) | ~ sP14), inference(cnf_transformation, [], [f100])).
fof(f902, plain, (~ spl30_95 | spl30_61), inference(avatar_split_clause, [], [f301, f641, f894])).
fof(f301, plain, ((e0 = op(e0, e0)) | ~ sP15), inference(cnf_transformation, [], [f99])).
fof(f99, plain, (((e3 = op(e0, op(e0, e3))) & (e2 = op(e0, op(e0, e2))) & (e1 = op(e0, op(e0, e1))) & (e0 = op(e0, op(e0, e0))) & (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP15), inference(nnf_transformation, [], [f69])).
fof(f900, plain, (~ spl30_95 | spl30_92), inference(avatar_split_clause, [], [f303, f867, f894])).
fof(f303, plain, ((e0 = op(e0, op(e0, e0))) | ~ sP15), inference(cnf_transformation, [], [f99])).
fof(f899, plain, (~ spl30_95 | spl30_91), inference(avatar_split_clause, [], [f304, f862, f894])).
fof(f304, plain, ((e1 = op(e0, op(e0, e1))) | ~ sP15), inference(cnf_transformation, [], [f99])).
fof(f898, plain, (~ spl30_95 | spl30_90), inference(avatar_split_clause, [], [f305, f857, f894])).
fof(f305, plain, ((e2 = op(e0, op(e0, e2))) | ~ sP15), inference(cnf_transformation, [], [f99])).
fof(f897, plain, (~ spl30_95 | spl30_89), inference(avatar_split_clause, [], [f306, f852, f894])).
fof(f306, plain, ((e3 = op(e0, op(e0, e3))) | ~ sP15), inference(cnf_transformation, [], [f99])).
fof(f892, plain, (~ spl30_94 | spl30_57), inference(avatar_split_clause, [], [f295, f624, f884])).
fof(f295, plain, ((e0 = op(e0, e1)) | ~ sP16), inference(cnf_transformation, [], [f98])).
fof(f98, plain, (((e3 = op(e0, op(e0, e3))) & (e2 = op(e0, op(e0, e2))) & (e1 = op(e0, op(e0, e1))) & (e0 = op(e0, op(e0, e0))) & (e0 = op(e1, e0)) & (e0 = op(e0, e1))) | ~ sP16), inference(nnf_transformation, [], [f70])).
fof(f891, plain, (~ spl30_94 | spl30_45), inference(avatar_split_clause, [], [f296, f573, f884])).
fof(f296, plain, ((e0 = op(e1, e0)) | ~ sP16), inference(cnf_transformation, [], [f98])).
fof(f890, plain, (~ spl30_94 | spl30_92), inference(avatar_split_clause, [], [f297, f867, f884])).
fof(f297, plain, ((e0 = op(e0, op(e0, e0))) | ~ sP16), inference(cnf_transformation, [], [f98])).
fof(f889, plain, (~ spl30_94 | spl30_91), inference(avatar_split_clause, [], [f298, f862, f884])).
fof(f298, plain, ((e1 = op(e0, op(e0, e1))) | ~ sP16), inference(cnf_transformation, [], [f98])).
fof(f888, plain, (~ spl30_94 | spl30_90), inference(avatar_split_clause, [], [f299, f857, f884])).
fof(f299, plain, ((e2 = op(e0, op(e0, e2))) | ~ sP16), inference(cnf_transformation, [], [f98])).
fof(f887, plain, (~ spl30_94 | spl30_89), inference(avatar_split_clause, [], [f300, f852, f884])).
fof(f300, plain, ((e3 = op(e0, op(e0, e3))) | ~ sP16), inference(cnf_transformation, [], [f98])).
fof(f882, plain, (~ spl30_93 | spl30_53), inference(avatar_split_clause, [], [f289, f607, f874])).
fof(f289, plain, ((e0 = op(e0, e2)) | ~ sP17), inference(cnf_transformation, [], [f97])).
fof(f97, plain, (((e3 = op(e0, op(e0, e3))) & (e2 = op(e0, op(e0, e2))) & (e1 = op(e0, op(e0, e1))) & (e0 = op(e0, op(e0, e0))) & (e0 = op(e2, e0)) & (e0 = op(e0, e2))) | ~ sP17), inference(nnf_transformation, [], [f71])).
fof(f881, plain, (~ spl30_93 | spl30_29), inference(avatar_split_clause, [], [f290, f505, f874])).
fof(f290, plain, ((e0 = op(e2, e0)) | ~ sP17), inference(cnf_transformation, [], [f97])).
fof(f880, plain, (~ spl30_93 | spl30_92), inference(avatar_split_clause, [], [f291, f867, f874])).
fof(f291, plain, ((e0 = op(e0, op(e0, e0))) | ~ sP17), inference(cnf_transformation, [], [f97])).
fof(f879, plain, (~ spl30_93 | spl30_91), inference(avatar_split_clause, [], [f292, f862, f874])).
fof(f292, plain, ((e1 = op(e0, op(e0, e1))) | ~ sP17), inference(cnf_transformation, [], [f97])).
fof(f878, plain, (~ spl30_93 | spl30_90), inference(avatar_split_clause, [], [f293, f857, f874])).
fof(f293, plain, ((e2 = op(e0, op(e0, e2))) | ~ sP17), inference(cnf_transformation, [], [f97])).
fof(f877, plain, (~ spl30_93 | spl30_89), inference(avatar_split_clause, [], [f294, f852, f874])).
fof(f294, plain, ((e3 = op(e0, op(e0, e3))) | ~ sP17), inference(cnf_transformation, [], [f97])).
fof(f872, plain, (~ spl30_88 | spl30_49), inference(avatar_split_clause, [], [f283, f590, f848])).
fof(f283, plain, ((e0 = op(e0, e3)) | ~ sP18), inference(cnf_transformation, [], [f96])).
fof(f96, plain, (((e3 = op(e0, op(e0, e3))) & (e2 = op(e0, op(e0, e2))) & (e1 = op(e0, op(e0, e1))) & (e0 = op(e0, op(e0, e0))) & (e0 = op(e3, e0)) & (e0 = op(e0, e3))) | ~ sP18), inference(nnf_transformation, [], [f72])).
fof(f871, plain, (~ spl30_88 | spl30_13), inference(avatar_split_clause, [], [f284, f437, f848])).
fof(f284, plain, ((e0 = op(e3, e0)) | ~ sP18), inference(cnf_transformation, [], [f96])).
fof(f870, plain, (~ spl30_88 | spl30_92), inference(avatar_split_clause, [], [f285, f867, f848])).
fof(f285, plain, ((e0 = op(e0, op(e0, e0))) | ~ sP18), inference(cnf_transformation, [], [f96])).
fof(f865, plain, (~ spl30_88 | spl30_91), inference(avatar_split_clause, [], [f286, f862, f848])).
fof(f286, plain, ((e1 = op(e0, op(e0, e1))) | ~ sP18), inference(cnf_transformation, [], [f96])).
fof(f860, plain, (~ spl30_88 | spl30_90), inference(avatar_split_clause, [], [f287, f857, f848])).
fof(f287, plain, ((e2 = op(e0, op(e0, e2))) | ~ sP18), inference(cnf_transformation, [], [f96])).
fof(f855, plain, (~ spl30_88 | spl30_89), inference(avatar_split_clause, [], [f288, f852, f848])).
fof(f288, plain, ((e3 = op(e0, op(e0, e3))) | ~ sP18), inference(cnf_transformation, [], [f96])).
fof(f846, plain, (~ spl30_87 | spl30_46), inference(avatar_split_clause, [], [f277, f577, f838])).
fof(f277, plain, ((e1 = op(e1, e0)) | ~ sP19), inference(cnf_transformation, [], [f95])).
fof(f95, plain, (((e3 = op(e1, op(e1, e3))) & (e2 = op(e1, op(e1, e2))) & (e1 = op(e1, op(e1, e1))) & (e0 = op(e1, op(e1, e0))) & (e1 = op(e0, e1)) & (e1 = op(e1, e0))) | ~ sP19), inference(nnf_transformation, [], [f73])).
fof(f845, plain, (~ spl30_87 | spl30_58), inference(avatar_split_clause, [], [f278, f628, f838])).
fof(f278, plain, ((e1 = op(e0, e1)) | ~ sP19), inference(cnf_transformation, [], [f95])).
fof(f844, plain, (~ spl30_87 | spl30_84), inference(avatar_split_clause, [], [f279, f811, f838])).
fof(f279, plain, ((e0 = op(e1, op(e1, e0))) | ~ sP19), inference(cnf_transformation, [], [f95])).
fof(f842, plain, (~ spl30_87 | spl30_82), inference(avatar_split_clause, [], [f281, f801, f838])).
fof(f281, plain, ((e2 = op(e1, op(e1, e2))) | ~ sP19), inference(cnf_transformation, [], [f95])).
fof(f841, plain, (~ spl30_87 | spl30_81), inference(avatar_split_clause, [], [f282, f796, f838])).
fof(f282, plain, ((e3 = op(e1, op(e1, e3))) | ~ sP19), inference(cnf_transformation, [], [f95])).
fof(f836, plain, (~ spl30_86 | spl30_42), inference(avatar_split_clause, [], [f271, f560, f828])).
fof(f271, plain, ((e1 = op(e1, e1)) | ~ sP20), inference(cnf_transformation, [], [f94])).
fof(f94, plain, (((e3 = op(e1, op(e1, e3))) & (e2 = op(e1, op(e1, e2))) & (e1 = op(e1, op(e1, e1))) & (e0 = op(e1, op(e1, e0))) & (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP20), inference(nnf_transformation, [], [f74])).
fof(f834, plain, (~ spl30_86 | spl30_84), inference(avatar_split_clause, [], [f273, f811, f828])).
fof(f273, plain, ((e0 = op(e1, op(e1, e0))) | ~ sP20), inference(cnf_transformation, [], [f94])).
fof(f833, plain, (~ spl30_86 | spl30_83), inference(avatar_split_clause, [], [f274, f806, f828])).
fof(f274, plain, ((e1 = op(e1, op(e1, e1))) | ~ sP20), inference(cnf_transformation, [], [f94])).
fof(f832, plain, (~ spl30_86 | spl30_82), inference(avatar_split_clause, [], [f275, f801, f828])).
fof(f275, plain, ((e2 = op(e1, op(e1, e2))) | ~ sP20), inference(cnf_transformation, [], [f94])).
fof(f831, plain, (~ spl30_86 | spl30_81), inference(avatar_split_clause, [], [f276, f796, f828])).
fof(f276, plain, ((e3 = op(e1, op(e1, e3))) | ~ sP20), inference(cnf_transformation, [], [f94])).
fof(f826, plain, (~ spl30_85 | spl30_38), inference(avatar_split_clause, [], [f265, f543, f818])).
fof(f265, plain, ((e1 = op(e1, e2)) | ~ sP21), inference(cnf_transformation, [], [f93])).
fof(f93, plain, (((e3 = op(e1, op(e1, e3))) & (e2 = op(e1, op(e1, e2))) & (e1 = op(e1, op(e1, e1))) & (e0 = op(e1, op(e1, e0))) & (e1 = op(e2, e1)) & (e1 = op(e1, e2))) | ~ sP21), inference(nnf_transformation, [], [f75])).
fof(f825, plain, (~ spl30_85 | spl30_26), inference(avatar_split_clause, [], [f266, f492, f818])).
fof(f266, plain, ((e1 = op(e2, e1)) | ~ sP21), inference(cnf_transformation, [], [f93])).
fof(f824, plain, (~ spl30_85 | spl30_84), inference(avatar_split_clause, [], [f267, f811, f818])).
fof(f267, plain, ((e0 = op(e1, op(e1, e0))) | ~ sP21), inference(cnf_transformation, [], [f93])).
fof(f823, plain, (~ spl30_85 | spl30_83), inference(avatar_split_clause, [], [f268, f806, f818])).
fof(f268, plain, ((e1 = op(e1, op(e1, e1))) | ~ sP21), inference(cnf_transformation, [], [f93])).
fof(f822, plain, (~ spl30_85 | spl30_82), inference(avatar_split_clause, [], [f269, f801, f818])).
fof(f269, plain, ((e2 = op(e1, op(e1, e2))) | ~ sP21), inference(cnf_transformation, [], [f93])).
fof(f821, plain, (~ spl30_85 | spl30_81), inference(avatar_split_clause, [], [f270, f796, f818])).
fof(f270, plain, ((e3 = op(e1, op(e1, e3))) | ~ sP21), inference(cnf_transformation, [], [f93])).
fof(f816, plain, (~ spl30_80 | spl30_34), inference(avatar_split_clause, [], [f259, f526, f792])).
fof(f259, plain, ((e1 = op(e1, e3)) | ~ sP22), inference(cnf_transformation, [], [f92])).
fof(f92, plain, (((e3 = op(e1, op(e1, e3))) & (e2 = op(e1, op(e1, e2))) & (e1 = op(e1, op(e1, e1))) & (e0 = op(e1, op(e1, e0))) & (e1 = op(e3, e1)) & (e1 = op(e1, e3))) | ~ sP22), inference(nnf_transformation, [], [f76])).
fof(f815, plain, (~ spl30_80 | spl30_10), inference(avatar_split_clause, [], [f260, f424, f792])).
fof(f260, plain, ((e1 = op(e3, e1)) | ~ sP22), inference(cnf_transformation, [], [f92])).
fof(f814, plain, (~ spl30_80 | spl30_84), inference(avatar_split_clause, [], [f261, f811, f792])).
fof(f261, plain, ((e0 = op(e1, op(e1, e0))) | ~ sP22), inference(cnf_transformation, [], [f92])).
fof(f809, plain, (~ spl30_80 | spl30_83), inference(avatar_split_clause, [], [f262, f806, f792])).
fof(f262, plain, ((e1 = op(e1, op(e1, e1))) | ~ sP22), inference(cnf_transformation, [], [f92])).
fof(f804, plain, (~ spl30_80 | spl30_82), inference(avatar_split_clause, [], [f263, f801, f792])).
fof(f263, plain, ((e2 = op(e1, op(e1, e2))) | ~ sP22), inference(cnf_transformation, [], [f92])).
fof(f799, plain, (~ spl30_80 | spl30_81), inference(avatar_split_clause, [], [f264, f796, f792])).
fof(f264, plain, ((e3 = op(e1, op(e1, e3))) | ~ sP22), inference(cnf_transformation, [], [f92])).
fof(f790, plain, (~ spl30_79 | spl30_31), inference(avatar_split_clause, [], [f253, f513, f782])).
fof(f253, plain, ((e2 = op(e2, e0)) | ~ sP23), inference(cnf_transformation, [], [f91])).
fof(f91, plain, (((e3 = op(e2, op(e2, e3))) & (e2 = op(e2, op(e2, e2))) & (e1 = op(e2, op(e2, e1))) & (e0 = op(e2, op(e2, e0))) & (e2 = op(e0, e2)) & (e2 = op(e2, e0))) | ~ sP23), inference(nnf_transformation, [], [f77])).
fof(f789, plain, (~ spl30_79 | spl30_55), inference(avatar_split_clause, [], [f254, f615, f782])).
fof(f254, plain, ((e2 = op(e0, e2)) | ~ sP23), inference(cnf_transformation, [], [f91])).
fof(f788, plain, (~ spl30_79 | spl30_76), inference(avatar_split_clause, [], [f255, f755, f782])).
fof(f255, plain, ((e0 = op(e2, op(e2, e0))) | ~ sP23), inference(cnf_transformation, [], [f91])).
fof(f787, plain, (~ spl30_79 | spl30_75), inference(avatar_split_clause, [], [f256, f750, f782])).
fof(f256, plain, ((e1 = op(e2, op(e2, e1))) | ~ sP23), inference(cnf_transformation, [], [f91])).
fof(f786, plain, (~ spl30_79 | spl30_74), inference(avatar_split_clause, [], [f257, f745, f782])).
fof(f257, plain, ((e2 = op(e2, op(e2, e2))) | ~ sP23), inference(cnf_transformation, [], [f91])).
fof(f785, plain, (~ spl30_79 | spl30_73), inference(avatar_split_clause, [], [f258, f740, f782])).
fof(f258, plain, ((e3 = op(e2, op(e2, e3))) | ~ sP23), inference(cnf_transformation, [], [f91])).
fof(f780, plain, (~ spl30_78 | spl30_27), inference(avatar_split_clause, [], [f247, f496, f772])).
fof(f247, plain, ((e2 = op(e2, e1)) | ~ sP24), inference(cnf_transformation, [], [f90])).
fof(f90, plain, (((e3 = op(e2, op(e2, e3))) & (e2 = op(e2, op(e2, e2))) & (e1 = op(e2, op(e2, e1))) & (e0 = op(e2, op(e2, e0))) & (e2 = op(e1, e2)) & (e2 = op(e2, e1))) | ~ sP24), inference(nnf_transformation, [], [f78])).
fof(f779, plain, (~ spl30_78 | spl30_39), inference(avatar_split_clause, [], [f248, f547, f772])).
fof(f248, plain, ((e2 = op(e1, e2)) | ~ sP24), inference(cnf_transformation, [], [f90])).
fof(f778, plain, (~ spl30_78 | spl30_76), inference(avatar_split_clause, [], [f249, f755, f772])).
fof(f249, plain, ((e0 = op(e2, op(e2, e0))) | ~ sP24), inference(cnf_transformation, [], [f90])).
fof(f775, plain, (~ spl30_78 | spl30_73), inference(avatar_split_clause, [], [f252, f740, f772])).
fof(f252, plain, ((e3 = op(e2, op(e2, e3))) | ~ sP24), inference(cnf_transformation, [], [f90])).
fof(f770, plain, (~ spl30_77 | spl30_23), inference(avatar_split_clause, [], [f241, f479, f762])).
fof(f241, plain, ((e2 = op(e2, e2)) | ~ sP25), inference(cnf_transformation, [], [f89])).
fof(f89, plain, (((e3 = op(e2, op(e2, e3))) & (e2 = op(e2, op(e2, e2))) & (e1 = op(e2, op(e2, e1))) & (e0 = op(e2, op(e2, e0))) & (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP25), inference(nnf_transformation, [], [f79])).
fof(f768, plain, (~ spl30_77 | spl30_76), inference(avatar_split_clause, [], [f243, f755, f762])).
fof(f243, plain, ((e0 = op(e2, op(e2, e0))) | ~ sP25), inference(cnf_transformation, [], [f89])).
fof(f767, plain, (~ spl30_77 | spl30_75), inference(avatar_split_clause, [], [f244, f750, f762])).
fof(f244, plain, ((e1 = op(e2, op(e2, e1))) | ~ sP25), inference(cnf_transformation, [], [f89])).
fof(f765, plain, (~ spl30_77 | spl30_73), inference(avatar_split_clause, [], [f246, f740, f762])).
fof(f246, plain, ((e3 = op(e2, op(e2, e3))) | ~ sP25), inference(cnf_transformation, [], [f89])).
fof(f760, plain, (~ spl30_72 | spl30_19), inference(avatar_split_clause, [], [f235, f462, f736])).
fof(f235, plain, ((e2 = op(e2, e3)) | ~ sP26), inference(cnf_transformation, [], [f88])).
fof(f88, plain, (((e3 = op(e2, op(e2, e3))) & (e2 = op(e2, op(e2, e2))) & (e1 = op(e2, op(e2, e1))) & (e0 = op(e2, op(e2, e0))) & (e2 = op(e3, e2)) & (e2 = op(e2, e3))) | ~ sP26), inference(nnf_transformation, [], [f80])).
fof(f759, plain, (~ spl30_72 | spl30_7), inference(avatar_split_clause, [], [f236, f411, f736])).
fof(f236, plain, ((e2 = op(e3, e2)) | ~ sP26), inference(cnf_transformation, [], [f88])).
fof(f758, plain, (~ spl30_72 | spl30_76), inference(avatar_split_clause, [], [f237, f755, f736])).
fof(f237, plain, ((e0 = op(e2, op(e2, e0))) | ~ sP26), inference(cnf_transformation, [], [f88])).
fof(f753, plain, (~ spl30_72 | spl30_75), inference(avatar_split_clause, [], [f238, f750, f736])).
fof(f238, plain, ((e1 = op(e2, op(e2, e1))) | ~ sP26), inference(cnf_transformation, [], [f88])).
fof(f743, plain, (~ spl30_72 | spl30_73), inference(avatar_split_clause, [], [f240, f740, f736])).
fof(f240, plain, ((e3 = op(e2, op(e2, e3))) | ~ sP26), inference(cnf_transformation, [], [f88])).
fof(f734, plain, (~ spl30_71 | spl30_16), inference(avatar_split_clause, [], [f229, f449, f726])).
fof(f229, plain, ((e3 = op(e3, e0)) | ~ sP27), inference(cnf_transformation, [], [f87])).
fof(f87, plain, (((e3 = op(e3, op(e3, e3))) & (e2 = op(e3, op(e3, e2))) & (e1 = op(e3, op(e3, e1))) & (e0 = op(e3, op(e3, e0))) & (e3 = op(e0, e3)) & (e3 = op(e3, e0))) | ~ sP27), inference(nnf_transformation, [], [f81])).
fof(f733, plain, (~ spl30_71 | spl30_52), inference(avatar_split_clause, [], [f230, f602, f726])).
fof(f230, plain, ((e3 = op(e0, e3)) | ~ sP27), inference(cnf_transformation, [], [f87])).
fof(f731, plain, (~ spl30_71 | spl30_68), inference(avatar_split_clause, [], [f232, f704, f726])).
fof(f232, plain, ((e1 = op(e3, op(e3, e1))) | ~ sP27), inference(cnf_transformation, [], [f87])).
fof(f730, plain, (~ spl30_71 | spl30_67), inference(avatar_split_clause, [], [f233, f699, f726])).
fof(f233, plain, ((e2 = op(e3, op(e3, e2))) | ~ sP27), inference(cnf_transformation, [], [f87])).
fof(f729, plain, (~ spl30_71 | spl30_66), inference(avatar_split_clause, [], [f234, f694, f726])).
fof(f234, plain, ((e3 = op(e3, op(e3, e3))) | ~ sP27), inference(cnf_transformation, [], [f87])).
fof(f724, plain, (~ spl30_70 | spl30_12), inference(avatar_split_clause, [], [f223, f432, f716])).
fof(f223, plain, ((e3 = op(e3, e1)) | ~ sP28), inference(cnf_transformation, [], [f86])).
fof(f86, plain, (((e3 = op(e3, op(e3, e3))) & (e2 = op(e3, op(e3, e2))) & (e1 = op(e3, op(e3, e1))) & (e0 = op(e3, op(e3, e0))) & (e3 = op(e1, e3)) & (e3 = op(e3, e1))) | ~ sP28), inference(nnf_transformation, [], [f82])).
fof(f723, plain, (~ spl30_70 | spl30_36), inference(avatar_split_clause, [], [f224, f534, f716])).
fof(f224, plain, ((e3 = op(e1, e3)) | ~ sP28), inference(cnf_transformation, [], [f86])).
fof(f722, plain, (~ spl30_70 | spl30_69), inference(avatar_split_clause, [], [f225, f709, f716])).
fof(f225, plain, ((e0 = op(e3, op(e3, e0))) | ~ sP28), inference(cnf_transformation, [], [f86])).
fof(f721, plain, (~ spl30_70 | spl30_68), inference(avatar_split_clause, [], [f226, f704, f716])).
fof(f226, plain, ((e1 = op(e3, op(e3, e1))) | ~ sP28), inference(cnf_transformation, [], [f86])).
fof(f720, plain, (~ spl30_70 | spl30_67), inference(avatar_split_clause, [], [f227, f699, f716])).
fof(f227, plain, ((e2 = op(e3, op(e3, e2))) | ~ sP28), inference(cnf_transformation, [], [f86])).
fof(f719, plain, (~ spl30_70 | spl30_66), inference(avatar_split_clause, [], [f228, f694, f716])).
fof(f228, plain, ((e3 = op(e3, op(e3, e3))) | ~ sP28), inference(cnf_transformation, [], [f86])).
fof(f714, plain, (~ spl30_65 | spl30_8), inference(avatar_split_clause, [], [f217, f415, f690])).
fof(f217, plain, ((e3 = op(e3, e2)) | ~ sP29), inference(cnf_transformation, [], [f85])).
fof(f85, plain, (((e3 = op(e3, op(e3, e3))) & (e2 = op(e3, op(e3, e2))) & (e1 = op(e3, op(e3, e1))) & (e0 = op(e3, op(e3, e0))) & (e3 = op(e2, e3)) & (e3 = op(e3, e2))) | ~ sP29), inference(nnf_transformation, [], [f83])).
fof(f713, plain, (~ spl30_65 | spl30_20), inference(avatar_split_clause, [], [f218, f466, f690])).
fof(f218, plain, ((e3 = op(e2, e3)) | ~ sP29), inference(cnf_transformation, [], [f85])).
fof(f712, plain, (~ spl30_65 | spl30_69), inference(avatar_split_clause, [], [f219, f709, f690])).
fof(f219, plain, ((e0 = op(e3, op(e3, e0))) | ~ sP29), inference(cnf_transformation, [], [f85])).
fof(f707, plain, (~ spl30_65 | spl30_68), inference(avatar_split_clause, [], [f220, f704, f690])).
fof(f220, plain, ((e1 = op(e3, op(e3, e1))) | ~ sP29), inference(cnf_transformation, [], [f85])).
fof(f702, plain, (~ spl30_65 | spl30_67), inference(avatar_split_clause, [], [f221, f699, f690])).
fof(f221, plain, ((e2 = op(e3, op(e3, e2))) | ~ sP29), inference(cnf_transformation, [], [f85])).
fof(f688, plain, (spl30_61 | spl30_57 | spl30_53 | spl30_49), inference(avatar_split_clause, [], [f131, f590, f607, f624, f641])).
fof(f131, plain, ((e0 = op(e0, e3)) | (e0 = op(e0, e2)) | (e0 = op(e0, e1)) | (e0 = op(e0, e0))), inference(cnf_transformation, [], [f2])).
fof(f2, plain, (((e3 = op(e3, e3)) | (e3 = op(e2, e3)) | (e3 = op(e1, e3)) | (e3 = op(e0, e3))) & ((e3 = op(e3, e3)) | (e3 = op(e3, e2)) | (e3 = op(e3, e1)) | (e3 = op(e3, e0))) & ((e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(e0, e3))) & ((e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, e0))) & ((e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(e0, e3))) & ((e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, e0))) & ((e0 = op(e3, e3)) | (e0 = op(e2, e3)) | (e0 = op(e1, e3)) | (e0 = op(e0, e3))) & ((e0 = op(e3, e3)) | (e0 = op(e3, e2)) | (e0 = op(e3, e1)) | (e0 = op(e3, e0))) & ((e3 = op(e3, e2)) | (e3 = op(e2, e2)) | (e3 = op(e1, e2)) | (e3 = op(e0, e2))) & ((e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, e0))) & ((e2 = op(e3, e2)) | (e2 = op(e2, e2)) | (e2 = op(e1, e2)) | (e2 = op(e0, e2))) & ((e2 = op(e2, e3)) | (e2 = op(e2, e2)) | (e2 = op(e2, e1)) | (e2 = op(e2, e0))) & ((e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(e0, e2))) & ((e1 = op(e2, e3)) | (e1 = op(e2, e2)) | (e1 = op(e2, e1)) | (e1 = op(e2, e0))) & ((e0 = op(e3, e2)) | (e0 = op(e2, e2)) | (e0 = op(e1, e2)) | (e0 = op(e0, e2))) & ((e0 = op(e2, e3)) | (e0 = op(e2, e2)) | (e0 = op(e2, e1)) | (e0 = op(e2, e0))) & ((e3 = op(e3, e1)) | (e3 = op(e2, e1)) | (e3 = op(e1, e1)) | (e3 = op(e0, e1))) & ((e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, e0))) & ((e2 = op(e3, e1)) | (e2 = op(e2, e1)) | (e2 = op(e1, e1)) | (e2 = op(e0, e1))) & ((e2 = op(e1, e3)) | (e2 = op(e1, e2)) | (e2 = op(e1, e1)) | (e2 = op(e1, e0))) & ((e1 = op(e3, e1)) | (e1 = op(e2, e1)) | (e1 = op(e1, e1)) | (e1 = op(e0, e1))) & ((e1 = op(e1, e3)) | (e1 = op(e1, e2)) | (e1 = op(e1, e1)) | (e1 = op(e1, e0))) & ((e0 = op(e3, e1)) | (e0 = op(e2, e1)) | (e0 = op(e1, e1)) | (e0 = op(e0, e1))) & ((e0 = op(e1, e3)) | (e0 = op(e1, e2)) | (e0 = op(e1, e1)) | (e0 = op(e1, e0))) & ((e3 = op(e3, e0)) | (e3 = op(e2, e0)) | (e3 = op(e1, e0)) | (op(e0, e0) = e3)) & ((e3 = op(e0, e3)) | (e3 = op(e0, e2)) | (e3 = op(e0, e1)) | (op(e0, e0) = e3)) & ((e2 = op(e3, e0)) | (e2 = op(e2, e0)) | (e2 = op(e1, e0)) | (op(e0, e0) = e2)) & ((e2 = op(e0, e3)) | (e2 = op(e0, e2)) | (e2 = op(e0, e1)) | (op(e0, e0) = e2)) & ((e1 = op(e3, e0)) | (e1 = op(e2, e0)) | (e1 = op(e1, e0)) | (op(e0, e0) = e1)) & ((e1 = op(e0, e3)) | (e1 = op(e0, e2)) | (e1 = op(e0, e1)) | (op(e0, e0) = e1)) & ((e0 = op(e3, e0)) | (e0 = op(e2, e0)) | (e0 = op(e1, e0)) | (e0 = op(e0, e0))) & ((e0 = op(e0, e3)) | (e0 = op(e0, e2)) | (e0 = op(e0, e1)) | (e0 = op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG164+1.p', ax2)).
fof(f687, plain, (spl30_61 | spl30_45 | spl30_29 | spl30_13), inference(avatar_split_clause, [], [f132, f437, f505, f573, f641])).
fof(f132, plain, ((e0 = op(e3, e0)) | (e0 = op(e2, e0)) | (e0 = op(e1, e0)) | (e0 = op(e0, e0))), inference(cnf_transformation, [], [f2])).
fof(f686, plain, (spl30_62 | spl30_58 | spl30_54 | spl30_50), inference(avatar_split_clause, [], [f133, f594, f611, f628, f645])).
fof(f133, plain, ((e1 = op(e0, e3)) | (e1 = op(e0, e2)) | (e1 = op(e0, e1)) | (op(e0, e0) = e1)), inference(cnf_transformation, [], [f2])).
fof(f685, plain, (spl30_62 | spl30_46 | spl30_30 | spl30_14), inference(avatar_split_clause, [], [f134, f441, f509, f577, f645])).
fof(f134, plain, ((e1 = op(e3, e0)) | (e1 = op(e2, e0)) | (e1 = op(e1, e0)) | (op(e0, e0) = e1)), inference(cnf_transformation, [], [f2])).
fof(f684, plain, (spl30_63 | spl30_59 | spl30_55 | spl30_51), inference(avatar_split_clause, [], [f135, f598, f615, f632, f649])).
fof(f135, plain, ((e2 = op(e0, e3)) | (e2 = op(e0, e2)) | (e2 = op(e0, e1)) | (op(e0, e0) = e2)), inference(cnf_transformation, [], [f2])).
fof(f683, plain, (spl30_63 | spl30_47 | spl30_31 | spl30_15), inference(avatar_split_clause, [], [f136, f445, f513, f581, f649])).
fof(f136, plain, ((e2 = op(e3, e0)) | (e2 = op(e2, e0)) | (e2 = op(e1, e0)) | (op(e0, e0) = e2)), inference(cnf_transformation, [], [f2])).
fof(f682, plain, (spl30_64 | spl30_60 | spl30_56 | spl30_52), inference(avatar_split_clause, [], [f137, f602, f619, f636, f653])).
fof(f137, plain, ((e3 = op(e0, e3)) | (e3 = op(e0, e2)) | (e3 = op(e0, e1)) | (op(e0, e0) = e3)), inference(cnf_transformation, [], [f2])).
fof(f681, plain, (spl30_64 | spl30_48 | spl30_32 | spl30_16), inference(avatar_split_clause, [], [f138, f449, f517, f585, f653])).
fof(f138, plain, ((e3 = op(e3, e0)) | (e3 = op(e2, e0)) | (e3 = op(e1, e0)) | (op(e0, e0) = e3)), inference(cnf_transformation, [], [f2])).
fof(f680, plain, (spl30_45 | spl30_41 | spl30_37 | spl30_33), inference(avatar_split_clause, [], [f139, f522, f539, f556, f573])).
fof(f139, plain, ((e0 = op(e1, e3)) | (e0 = op(e1, e2)) | (e0 = op(e1, e1)) | (e0 = op(e1, e0))), inference(cnf_transformation, [], [f2])).
fof(f679, plain, (spl30_57 | spl30_41 | spl30_25 | spl30_9), inference(avatar_split_clause, [], [f140, f420, f488, f556, f624])).
fof(f140, plain, ((e0 = op(e3, e1)) | (e0 = op(e2, e1)) | (e0 = op(e1, e1)) | (e0 = op(e0, e1))), inference(cnf_transformation, [], [f2])).
fof(f678, plain, (spl30_46 | spl30_42 | spl30_38 | spl30_34), inference(avatar_split_clause, [], [f141, f526, f543, f560, f577])).
fof(f141, plain, ((e1 = op(e1, e3)) | (e1 = op(e1, e2)) | (e1 = op(e1, e1)) | (e1 = op(e1, e0))), inference(cnf_transformation, [], [f2])).
fof(f677, plain, (spl30_58 | spl30_42 | spl30_26 | spl30_10), inference(avatar_split_clause, [], [f142, f424, f492, f560, f628])).
fof(f142, plain, ((e1 = op(e3, e1)) | (e1 = op(e2, e1)) | (e1 = op(e1, e1)) | (e1 = op(e0, e1))), inference(cnf_transformation, [], [f2])).
fof(f676, plain, (spl30_47 | spl30_43 | spl30_39 | spl30_35), inference(avatar_split_clause, [], [f143, f530, f547, f564, f581])).
fof(f143, plain, ((e2 = op(e1, e3)) | (e2 = op(e1, e2)) | (e2 = op(e1, e1)) | (e2 = op(e1, e0))), inference(cnf_transformation, [], [f2])).
fof(f675, plain, (spl30_59 | spl30_43 | spl30_27 | spl30_11), inference(avatar_split_clause, [], [f144, f428, f496, f564, f632])).
fof(f144, plain, ((e2 = op(e3, e1)) | (e2 = op(e2, e1)) | (e2 = op(e1, e1)) | (e2 = op(e0, e1))), inference(cnf_transformation, [], [f2])).
fof(f674, plain, (spl30_48 | spl30_44 | spl30_40 | spl30_36), inference(avatar_split_clause, [], [f145, f534, f551, f568, f585])).
fof(f145, plain, ((e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, e0))), inference(cnf_transformation, [], [f2])).
fof(f673, plain, (spl30_60 | spl30_44 | spl30_28 | spl30_12), inference(avatar_split_clause, [], [f146, f432, f500, f568, f636])).
fof(f146, plain, ((e3 = op(e3, e1)) | (e3 = op(e2, e1)) | (e3 = op(e1, e1)) | (e3 = op(e0, e1))), inference(cnf_transformation, [], [f2])).
fof(f672, plain, (spl30_29 | spl30_25 | spl30_21 | spl30_17), inference(avatar_split_clause, [], [f147, f454, f471, f488, f505])).
fof(f147, plain, ((e0 = op(e2, e3)) | (e0 = op(e2, e2)) | (e0 = op(e2, e1)) | (e0 = op(e2, e0))), inference(cnf_transformation, [], [f2])).
fof(f671, plain, (spl30_53 | spl30_37 | spl30_21 | spl30_5), inference(avatar_split_clause, [], [f148, f403, f471, f539, f607])).
fof(f148, plain, ((e0 = op(e3, e2)) | (e0 = op(e2, e2)) | (e0 = op(e1, e2)) | (e0 = op(e0, e2))), inference(cnf_transformation, [], [f2])).
fof(f670, plain, (spl30_30 | spl30_26 | spl30_22 | spl30_18), inference(avatar_split_clause, [], [f149, f458, f475, f492, f509])).
fof(f149, plain, ((e1 = op(e2, e3)) | (e1 = op(e2, e2)) | (e1 = op(e2, e1)) | (e1 = op(e2, e0))), inference(cnf_transformation, [], [f2])).
fof(f669, plain, (spl30_54 | spl30_38 | spl30_22 | spl30_6), inference(avatar_split_clause, [], [f150, f407, f475, f543, f611])).
fof(f150, plain, ((e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(e0, e2))), inference(cnf_transformation, [], [f2])).
fof(f668, plain, (spl30_31 | spl30_27 | spl30_23 | spl30_19), inference(avatar_split_clause, [], [f151, f462, f479, f496, f513])).
fof(f151, plain, ((e2 = op(e2, e3)) | (e2 = op(e2, e2)) | (e2 = op(e2, e1)) | (e2 = op(e2, e0))), inference(cnf_transformation, [], [f2])).
fof(f667, plain, (spl30_55 | spl30_39 | spl30_23 | spl30_7), inference(avatar_split_clause, [], [f152, f411, f479, f547, f615])).
fof(f152, plain, ((e2 = op(e3, e2)) | (e2 = op(e2, e2)) | (e2 = op(e1, e2)) | (e2 = op(e0, e2))), inference(cnf_transformation, [], [f2])).
fof(f666, plain, (spl30_32 | spl30_28 | spl30_24 | spl30_20), inference(avatar_split_clause, [], [f153, f466, f483, f500, f517])).
fof(f153, plain, ((e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, e0))), inference(cnf_transformation, [], [f2])).
fof(f665, plain, (spl30_56 | spl30_40 | spl30_24 | spl30_8), inference(avatar_split_clause, [], [f154, f415, f483, f551, f619])).
fof(f154, plain, ((e3 = op(e3, e2)) | (e3 = op(e2, e2)) | (e3 = op(e1, e2)) | (e3 = op(e0, e2))), inference(cnf_transformation, [], [f2])).
fof(f664, plain, (spl30_13 | spl30_9 | spl30_5 | spl30_1), inference(avatar_split_clause, [], [f155, f386, f403, f420, f437])).
fof(f155, plain, ((e0 = op(e3, e3)) | (e0 = op(e3, e2)) | (e0 = op(e3, e1)) | (e0 = op(e3, e0))), inference(cnf_transformation, [], [f2])).
fof(f663, plain, (spl30_49 | spl30_33 | spl30_17 | spl30_1), inference(avatar_split_clause, [], [f156, f386, f454, f522, f590])).
fof(f156, plain, ((e0 = op(e3, e3)) | (e0 = op(e2, e3)) | (e0 = op(e1, e3)) | (e0 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f662, plain, (spl30_14 | spl30_10 | spl30_6 | spl30_2), inference(avatar_split_clause, [], [f157, f390, f407, f424, f441])).
fof(f157, plain, ((e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, e0))), inference(cnf_transformation, [], [f2])).
fof(f661, plain, (spl30_50 | spl30_34 | spl30_18 | spl30_2), inference(avatar_split_clause, [], [f158, f390, f458, f526, f594])).
fof(f158, plain, ((e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f660, plain, (spl30_15 | spl30_11 | spl30_7 | spl30_3), inference(avatar_split_clause, [], [f159, f394, f411, f428, f445])).
fof(f159, plain, ((e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, e0))), inference(cnf_transformation, [], [f2])).
fof(f659, plain, (spl30_51 | spl30_35 | spl30_19 | spl30_3), inference(avatar_split_clause, [], [f160, f394, f462, f530, f598])).
fof(f160, plain, ((e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f658, plain, (spl30_16 | spl30_12 | spl30_8 | spl30_4), inference(avatar_split_clause, [], [f161, f398, f415, f432, f449])).
fof(f161, plain, ((e3 = op(e3, e3)) | (e3 = op(e3, e2)) | (e3 = op(e3, e1)) | (e3 = op(e3, e0))), inference(cnf_transformation, [], [f2])).
fof(f657, plain, (spl30_52 | spl30_36 | spl30_20 | spl30_4), inference(avatar_split_clause, [], [f162, f398, f466, f534, f602])).
fof(f162, plain, ((e3 = op(e3, e3)) | (e3 = op(e2, e3)) | (e3 = op(e1, e3)) | (e3 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f656, plain, (spl30_61 | spl30_62 | spl30_63 | spl30_64), inference(avatar_split_clause, [], [f115, f653, f649, f645, f641])).
fof(f115, plain, ((op(e0, e0) = e3) | (op(e0, e0) = e2) | (op(e0, e0) = e1) | (e0 = op(e0, e0))), inference(cnf_transformation, [], [f1])).
fof(f1, plain, (((e3 = op(e3, e3)) | (e2 = op(e3, e3)) | (e1 = op(e3, e3)) | (e0 = op(e3, e3))) & ((e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (e0 = op(e3, e2))) & ((e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (e0 = op(e3, e1))) & ((e3 = op(e3, e0)) | (e2 = op(e3, e0)) | (e1 = op(e3, e0)) | (e0 = op(e3, e0))) & ((e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))) & ((e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (e0 = op(e2, e2))) & ((e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))) & ((e3 = op(e2, e0)) | (e2 = op(e2, e0)) | (e1 = op(e2, e0)) | (e0 = op(e2, e0))) & ((e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (e0 = op(e1, e3))) & ((e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (e0 = op(e1, e2))) & ((e3 = op(e1, e1)) | (e2 = op(e1, e1)) | (e1 = op(e1, e1)) | (e0 = op(e1, e1))) & ((e3 = op(e1, e0)) | (e2 = op(e1, e0)) | (e1 = op(e1, e0)) | (e0 = op(e1, e0))) & ((e3 = op(e0, e3)) | (e2 = op(e0, e3)) | (e1 = op(e0, e3)) | (e0 = op(e0, e3))) & ((e3 = op(e0, e2)) | (e2 = op(e0, e2)) | (e1 = op(e0, e2)) | (e0 = op(e0, e2))) & ((e3 = op(e0, e1)) | (e2 = op(e0, e1)) | (e1 = op(e0, e1)) | (e0 = op(e0, e1))) & ((op(e0, e0) = e3) | (op(e0, e0) = e2) | (op(e0, e0) = e1) | (e0 = op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG164+1.p', ax1)).
fof(f639, plain, (spl30_57 | spl30_58 | spl30_59 | spl30_60), inference(avatar_split_clause, [], [f116, f636, f632, f628, f624])).
fof(f116, plain, ((e3 = op(e0, e1)) | (e2 = op(e0, e1)) | (e1 = op(e0, e1)) | (e0 = op(e0, e1))), inference(cnf_transformation, [], [f1])).
fof(f622, plain, (spl30_53 | spl30_54 | spl30_55 | spl30_56), inference(avatar_split_clause, [], [f117, f619, f615, f611, f607])).
fof(f117, plain, ((e3 = op(e0, e2)) | (e2 = op(e0, e2)) | (e1 = op(e0, e2)) | (e0 = op(e0, e2))), inference(cnf_transformation, [], [f1])).
fof(f605, plain, (spl30_49 | spl30_50 | spl30_51 | spl30_52), inference(avatar_split_clause, [], [f118, f602, f598, f594, f590])).
fof(f118, plain, ((e3 = op(e0, e3)) | (e2 = op(e0, e3)) | (e1 = op(e0, e3)) | (e0 = op(e0, e3))), inference(cnf_transformation, [], [f1])).
fof(f588, plain, (spl30_45 | spl30_46 | spl30_47 | spl30_48), inference(avatar_split_clause, [], [f119, f585, f581, f577, f573])).
fof(f119, plain, ((e3 = op(e1, e0)) | (e2 = op(e1, e0)) | (e1 = op(e1, e0)) | (e0 = op(e1, e0))), inference(cnf_transformation, [], [f1])).
fof(f571, plain, (spl30_41 | spl30_42 | spl30_43 | spl30_44), inference(avatar_split_clause, [], [f120, f568, f564, f560, f556])).
fof(f120, plain, ((e3 = op(e1, e1)) | (e2 = op(e1, e1)) | (e1 = op(e1, e1)) | (e0 = op(e1, e1))), inference(cnf_transformation, [], [f1])).
fof(f554, plain, (spl30_37 | spl30_38 | spl30_39 | spl30_40), inference(avatar_split_clause, [], [f121, f551, f547, f543, f539])).
fof(f121, plain, ((e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (e0 = op(e1, e2))), inference(cnf_transformation, [], [f1])).
fof(f537, plain, (spl30_33 | spl30_34 | spl30_35 | spl30_36), inference(avatar_split_clause, [], [f122, f534, f530, f526, f522])).
fof(f122, plain, ((e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (e0 = op(e1, e3))), inference(cnf_transformation, [], [f1])).
fof(f520, plain, (spl30_29 | spl30_30 | spl30_31 | spl30_32), inference(avatar_split_clause, [], [f123, f517, f513, f509, f505])).
fof(f123, plain, ((e3 = op(e2, e0)) | (e2 = op(e2, e0)) | (e1 = op(e2, e0)) | (e0 = op(e2, e0))), inference(cnf_transformation, [], [f1])).
fof(f503, plain, (spl30_25 | spl30_26 | spl30_27 | spl30_28), inference(avatar_split_clause, [], [f124, f500, f496, f492, f488])).
fof(f124, plain, ((e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))), inference(cnf_transformation, [], [f1])).
fof(f486, plain, (spl30_21 | spl30_22 | spl30_23 | spl30_24), inference(avatar_split_clause, [], [f125, f483, f479, f475, f471])).
fof(f125, plain, ((e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (e0 = op(e2, e2))), inference(cnf_transformation, [], [f1])).
fof(f469, plain, (spl30_17 | spl30_18 | spl30_19 | spl30_20), inference(avatar_split_clause, [], [f126, f466, f462, f458, f454])).
fof(f126, plain, ((e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))), inference(cnf_transformation, [], [f1])).
fof(f452, plain, (spl30_13 | spl30_14 | spl30_15 | spl30_16), inference(avatar_split_clause, [], [f127, f449, f445, f441, f437])).
fof(f127, plain, ((e3 = op(e3, e0)) | (e2 = op(e3, e0)) | (e1 = op(e3, e0)) | (e0 = op(e3, e0))), inference(cnf_transformation, [], [f1])).
fof(f435, plain, (spl30_9 | spl30_10 | spl30_11 | spl30_12), inference(avatar_split_clause, [], [f128, f432, f428, f424, f420])).
fof(f128, plain, ((e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (e0 = op(e3, e1))), inference(cnf_transformation, [], [f1])).
fof(f418, plain, (spl30_5 | spl30_6 | spl30_7 | spl30_8), inference(avatar_split_clause, [], [f129, f415, f411, f407, f403])).
fof(f129, plain, ((e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (e0 = op(e3, e2))), inference(cnf_transformation, [], [f1])).
fof(f401, plain, (spl30_1 | spl30_2 | spl30_3 | spl30_4), inference(avatar_split_clause, [], [f130, f398, f394, f390, f386])).
fof(f130, plain, ((e3 = op(e3, e3)) | (e2 = op(e3, e3)) | (e1 = op(e3, e3)) | (e0 = op(e3, e3))), inference(cnf_transformation, [], [f1])).