fof(f3390, plain, $false, inference(avatar_sat_refutation, [], [f747, f819, f826, f828, f829, f831, f832, f833, f841, f842, f881, f932, f1051, f1068, f1085, f1123, f1128, f1129, f1135, f1140, f1146, f1262, f1264, f1266, f1269, f1270, f1271, f1272, f1381, f1383, f1384, f1391, f1392, f1395, f1396, f1398, f1399, f1447, f1465, f1486, f1693, f2108, f2111, f2125, f2203, f2206, f2229, f2239, f2251, f2253, f2259, f2277, f2300, f2305, f2336, f2338, f2341, f2357, f2364, f2383, f2385, f2389, f2391, f2405, f2432, f2447, f2456, f2465, f2472, f2479, f2537, f2540, f2565, f2578, f2582, f2584, f2640, f2647, f2679, f2697, f2705, f2706, f2711, f2763, f2764, f2768, f2773, f2795, f2798, f2803, f2807, f2823, f2873, f2881, f2892, f2899, f2933, f2980, f3021, f3054, f3078, f3111, f3142, f3178, f3208, f3239, f3264, f3284, f3306, f3351])).
fof(f3351, plain, (~ spl42_61 | ~ spl42_62), inference(avatar_contradiction_clause, [], [f3350])).
fof(f3350, plain, ($false | (~ spl42_61 | ~ spl42_62)), inference(subsumption_resolution, [], [f3349, f300])).
fof(f300, plain, ~ (e10 = e11), inference(cnf_transformation, [], [f7])).
fof(f7, plain, (~ (e12 = e13) & ~ (e11 = e13) & ~ (e11 = e12) & ~ (e10 = e13) & ~ (e10 = e12) & ~ (e10 = e11)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG121+1.p', ax7)).
fof(f3349, plain, ((e10 = e11) | (~ spl42_61 | ~ spl42_62)), inference(backward_demodulation, [], [f806, f802])).
fof(f802, plain, ((e10 = op1(e10, e10)) | ~ spl42_61), inference(avatar_component_clause, [], [f800])).
fof(f800, plain, (spl42_61 <=> (e10 = op1(e10, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_61])])).
fof(f806, plain, ((op1(e10, e10) = e11) | ~ spl42_62), inference(avatar_component_clause, [], [f804])).
fof(f804, plain, (spl42_62 <=> (op1(e10, e10) = e11)), introduced(avatar_definition, [new_symbols(naming, [spl42_62])])).
fof(f3306, plain, (~ spl42_57 | ~ spl42_121 | ~ spl42_167 | spl42_206), inference(avatar_contradiction_clause, [], [f3305])).
fof(f3305, plain, ($false | (~ spl42_57 | ~ spl42_121 | ~ spl42_167 | spl42_206)), inference(subsumption_resolution, [], [f3304, f1089])).
fof(f1089, plain, ((e20 = op2(e20, e21)) | ~ spl42_121), inference(avatar_component_clause, [], [f1087])).
fof(f1087, plain, (spl42_121 <=> (e20 = op2(e20, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_121])])).
fof(f3304, plain, (~ (e20 = op2(e20, e21)) | (~ spl42_57 | ~ spl42_167 | spl42_206)), inference(forward_demodulation, [], [f3303, f1484])).
fof(f1484, plain, ((e20 = h4(e10)) | ~ spl42_167), inference(avatar_component_clause, [], [f1483])).
fof(f1483, plain, (spl42_167 <=> (e20 = h4(e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_167])])).
fof(f3303, plain, (~ (op2(e20, e21) = h4(e10)) | (~ spl42_57 | ~ spl42_167 | spl42_206)), inference(forward_demodulation, [], [f3302, f785])).
fof(f785, plain, ((e10 = op1(e10, e11)) | ~ spl42_57), inference(avatar_component_clause, [], [f783])).
fof(f783, plain, (spl42_57 <=> (e10 = op1(e10, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_57])])).
fof(f3302, plain, (~ (op2(e20, e21) = h4(op1(e10, e11))) | (~ spl42_167 | spl42_206)), inference(forward_demodulation, [], [f1688, f1484])).
fof(f1688, plain, (~ (h4(op1(e10, e11)) = op2(h4(e10), e21)) | spl42_206), inference(avatar_component_clause, [], [f1686])).
fof(f1686, plain, (spl42_206 <=> (h4(op1(e10, e11)) = op2(h4(e10), e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_206])])).
fof(f3284, plain, (~ spl42_6 | ~ spl42_70 | ~ spl42_160 | spl42_196), inference(avatar_contradiction_clause, [], [f3283])).
fof(f3283, plain, ($false | (~ spl42_6 | ~ spl42_70 | ~ spl42_160 | spl42_196)), inference(subsumption_resolution, [], [f3282, f872])).
fof(f872, plain, ((e21 = op2(e23, e22)) | ~ spl42_70), inference(avatar_component_clause, [], [f870])).
fof(f870, plain, (spl42_70 <=> (e21 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_70])])).
fof(f3282, plain, (~ (e21 = op2(e23, e22)) | (~ spl42_6 | ~ spl42_160 | spl42_196)), inference(forward_demodulation, [], [f3281, f1430])).
fof(f1430, plain, (e21 = h4(e11)), inference(forward_demodulation, [], [f478, f462])).
fof(f462, plain, (e21 = op2(op2(e23, e23), op2(e23, e23))), inference(cnf_transformation, [], [f13])).
fof(f13, plain, ((e22 = op2(op2(op2(e23, e23), op2(e23, e23)), e23)) & (e21 = op2(op2(e23, e23), op2(e23, e23))) & (e20 = op2(e23, e23))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG121+1.p', ax13)).
fof(f478, plain, (op2(op2(e23, e23), op2(e23, e23)) = h4(e11)), inference(cnf_transformation, [], [f17])).
fof(f17, plain, ((op2(op2(op2(e23, e23), op2(e23, e23)), e23) = h4(e12)) & (op2(op2(e23, e23), op2(e23, e23)) = h4(e11)) & (op2(e23, e23) = h4(e10)) & (e23 = h4(e13))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG121+1.p', ax17)).
fof(f3281, plain, (~ (op2(e23, e22) = h4(e11)) | (~ spl42_6 | ~ spl42_160 | spl42_196)), inference(forward_demodulation, [], [f3280, f568])).
fof(f568, plain, ((e11 = op1(e13, e12)) | ~ spl42_6), inference(avatar_component_clause, [], [f566])).
fof(f566, plain, (spl42_6 <=> (e11 = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_6])])).
fof(f3280, plain, (~ (op2(e23, e22) = h4(op1(e13, e12))) | (~ spl42_160 | spl42_196)), inference(forward_demodulation, [], [f1648, f1445])).
fof(f1445, plain, ((e22 = h4(e12)) | ~ spl42_160), inference(avatar_component_clause, [], [f1444])).
fof(f1444, plain, (spl42_160 <=> (e22 = h4(e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_160])])).
fof(f1648, plain, (~ (h4(op1(e13, e12)) = op2(e23, h4(e12))) | spl42_196), inference(avatar_component_clause, [], [f1646])).
fof(f1646, plain, (spl42_196 <=> (h4(op1(e13, e12)) = op2(e23, h4(e12)))), introduced(avatar_definition, [new_symbols(naming, [spl42_196])])).
fof(f3264, plain, (~ spl42_28 | ~ spl42_92 | ~ spl42_160 | spl42_200), inference(avatar_contradiction_clause, [], [f3263])).
fof(f3263, plain, ($false | (~ spl42_28 | ~ spl42_92 | ~ spl42_160 | spl42_200)), inference(subsumption_resolution, [], [f3262, f965])).
fof(f965, plain, ((e23 = op2(e22, e21)) | ~ spl42_92), inference(avatar_component_clause, [], [f963])).
fof(f963, plain, (spl42_92 <=> (e23 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_92])])).
fof(f3262, plain, (~ (e23 = op2(e22, e21)) | (~ spl42_28 | ~ spl42_160 | spl42_200)), inference(forward_demodulation, [], [f3261, f476])).
fof(f476, plain, (e23 = h4(e13)), inference(cnf_transformation, [], [f17])).
fof(f3261, plain, (~ (op2(e22, e21) = h4(e13)) | (~ spl42_28 | ~ spl42_160 | spl42_200)), inference(forward_demodulation, [], [f3260, f661])).
fof(f661, plain, ((e13 = op1(e12, e11)) | ~ spl42_28), inference(avatar_component_clause, [], [f659])).
fof(f659, plain, (spl42_28 <=> (e13 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_28])])).
fof(f3260, plain, (~ (op2(e22, e21) = h4(op1(e12, e11))) | (~ spl42_160 | spl42_200)), inference(forward_demodulation, [], [f1664, f1445])).
fof(f1664, plain, (~ (h4(op1(e12, e11)) = op2(h4(e12), e21)) | spl42_200), inference(avatar_component_clause, [], [f1662])).
fof(f1662, plain, (spl42_200 <=> (h4(op1(e12, e11)) = op2(h4(e12), e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_200])])).
fof(f3239, plain, (~ spl42_62 | spl42_207), inference(avatar_contradiction_clause, [], [f3238])).
fof(f3238, plain, ($false | (~ spl42_62 | spl42_207)), inference(subsumption_resolution, [], [f3237, f1430])).
fof(f3237, plain, (~ (e21 = h4(e11)) | (~ spl42_62 | spl42_207)), inference(forward_demodulation, [], [f1692, f806])).
fof(f1692, plain, (~ (e21 = h4(op1(e10, e10))) | spl42_207), inference(avatar_component_clause, [], [f1690])).
fof(f1690, plain, (spl42_207 <=> (e21 = h4(op1(e10, e10)))), introduced(avatar_definition, [new_symbols(naming, [spl42_207])])).
fof(f3208, plain, (~ spl42_52 | ~ spl42_116 | ~ spl42_167 | spl42_205), inference(avatar_contradiction_clause, [], [f3207])).
fof(f3207, plain, ($false | (~ spl42_52 | ~ spl42_116 | ~ spl42_167 | spl42_205)), inference(subsumption_resolution, [], [f3206, f1067])).
fof(f1067, plain, ((e23 = op2(e20, e23)) | ~ spl42_116), inference(avatar_component_clause, [], [f1065])).
fof(f1065, plain, (spl42_116 <=> (e23 = op2(e20, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_116])])).
fof(f3206, plain, (~ (e23 = op2(e20, e23)) | (~ spl42_52 | ~ spl42_167 | spl42_205)), inference(forward_demodulation, [], [f3205, f476])).
fof(f3205, plain, (~ (op2(e20, e23) = h4(e13)) | (~ spl42_52 | ~ spl42_167 | spl42_205)), inference(forward_demodulation, [], [f3204, f763])).
fof(f763, plain, ((e13 = op1(e10, e13)) | ~ spl42_52), inference(avatar_component_clause, [], [f761])).
fof(f761, plain, (spl42_52 <=> (e13 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_52])])).
fof(f3204, plain, (~ (op2(e20, e23) = h4(op1(e10, e13))) | (~ spl42_167 | spl42_205)), inference(forward_demodulation, [], [f1684, f1484])).
fof(f1684, plain, (~ (h4(op1(e10, e13)) = op2(h4(e10), e23)) | spl42_205), inference(avatar_component_clause, [], [f1682])).
fof(f1682, plain, (spl42_205 <=> (h4(op1(e10, e13)) = op2(h4(e10), e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_205])])).
fof(f3178, plain, (~ spl42_11 | ~ spl42_75 | ~ spl42_160 | spl42_197), inference(avatar_contradiction_clause, [], [f3177])).
fof(f3177, plain, ($false | (~ spl42_11 | ~ spl42_75 | ~ spl42_160 | spl42_197)), inference(subsumption_resolution, [], [f3176, f893])).
fof(f893, plain, ((e22 = op2(e23, e21)) | ~ spl42_75), inference(avatar_component_clause, [], [f891])).
fof(f891, plain, (spl42_75 <=> (e22 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_75])])).
fof(f3176, plain, (~ (e22 = op2(e23, e21)) | (~ spl42_11 | ~ spl42_160 | spl42_197)), inference(forward_demodulation, [], [f3175, f1445])).
fof(f3175, plain, (~ (op2(e23, e21) = h4(e12)) | (~ spl42_11 | spl42_197)), inference(forward_demodulation, [], [f1652, f589])).
fof(f589, plain, ((e12 = op1(e13, e11)) | ~ spl42_11), inference(avatar_component_clause, [], [f587])).
fof(f587, plain, (spl42_11 <=> (e12 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_11])])).
fof(f1652, plain, (~ (op2(e23, e21) = h4(op1(e13, e11))) | spl42_197), inference(avatar_component_clause, [], [f1650])).
fof(f1650, plain, (spl42_197 <=> (op2(e23, e21) = h4(op1(e13, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl42_197])])).
fof(f3142, plain, (~ spl42_42 | ~ spl42_190 | ~ spl42_191 | spl42_203), inference(avatar_contradiction_clause, [], [f3141])).
fof(f3141, plain, ($false | (~ spl42_42 | ~ spl42_190 | ~ spl42_191 | spl42_203)), inference(subsumption_resolution, [], [f3140, f2937])).
fof(f2937, plain, ((e21 = h2(e10)) | (~ spl42_190 | ~ spl42_191)), inference(backward_demodulation, [], [f2828, f1601])).
fof(f1601, plain, ((e21 = h1(e11)) | ~ spl42_190), inference(avatar_component_clause, [], [f1600])).
fof(f1600, plain, (spl42_190 <=> (e21 = h1(e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_190])])).
fof(f2828, plain, ((h1(e11) = h2(e10)) | ~ spl42_191), inference(backward_demodulation, [], [f469, f2827])).
fof(f2827, plain, ((op2(e21, e21) = h1(e11)) | ~ spl42_191), inference(backward_demodulation, [], [f1407, f1606])).
fof(f1606, plain, ((e21 = h1(e10)) | ~ spl42_191), inference(avatar_component_clause, [], [f1605])).
fof(f1605, plain, (spl42_191 <=> (e21 = h1(e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_191])])).
fof(f1407, plain, (h1(e11) = op2(h1(e10), h1(e10))), inference(backward_demodulation, [], [f466, f465])).
fof(f465, plain, (op2(e20, e20) = h1(e10)), inference(cnf_transformation, [], [f14])).
fof(f14, plain, ((h1(e12) = op2(op2(op2(e20, e20), op2(e20, e20)), e20)) & (h1(e11) = op2(op2(e20, e20), op2(e20, e20))) & (op2(e20, e20) = h1(e10)) & (e20 = h1(e13))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG121+1.p', ax14)).
fof(f466, plain, (h1(e11) = op2(op2(e20, e20), op2(e20, e20))), inference(cnf_transformation, [], [f14])).
fof(f469, plain, (op2(e21, e21) = h2(e10)), inference(cnf_transformation, [], [f15])).
fof(f15, plain, ((h2(e12) = op2(op2(op2(e21, e21), op2(e21, e21)), e21)) & (h2(e11) = op2(op2(e21, e21), op2(e21, e21))) & (op2(e21, e21) = h2(e10)) & (e21 = h2(e13))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG121+1.p', ax15)).
fof(f3140, plain, (~ (e21 = h2(e10)) | (~ spl42_42 | spl42_203)), inference(forward_demodulation, [], [f3139, f1430])).
fof(f3139, plain, (~ (h2(e10) = h4(e11)) | (~ spl42_42 | spl42_203)), inference(forward_demodulation, [], [f1676, f721])).
fof(f721, plain, ((e11 = op1(e11, e11)) | ~ spl42_42), inference(avatar_component_clause, [], [f719])).
fof(f719, plain, (spl42_42 <=> (e11 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_42])])).
fof(f1676, plain, (~ (h2(e10) = h4(op1(e11, e11))) | spl42_203), inference(avatar_component_clause, [], [f1674])).
fof(f1674, plain, (spl42_203 <=> (h2(e10) = h4(op1(e11, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl42_203])])).
fof(f3111, plain, (~ spl42_40 | ~ spl42_104 | ~ spl42_160 | spl42_202), inference(avatar_contradiction_clause, [], [f3110])).
fof(f3110, plain, ($false | (~ spl42_40 | ~ spl42_104 | ~ spl42_160 | spl42_202)), inference(subsumption_resolution, [], [f3109, f1016])).
fof(f1016, plain, ((e23 = op2(e21, e22)) | ~ spl42_104), inference(avatar_component_clause, [], [f1014])).
fof(f1014, plain, (spl42_104 <=> (e23 = op2(e21, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_104])])).
fof(f3109, plain, (~ (e23 = op2(e21, e22)) | (~ spl42_40 | ~ spl42_160 | spl42_202)), inference(forward_demodulation, [], [f3108, f476])).
fof(f3108, plain, (~ (op2(e21, e22) = h4(e13)) | (~ spl42_40 | ~ spl42_160 | spl42_202)), inference(forward_demodulation, [], [f3107, f712])).
fof(f712, plain, ((e13 = op1(e11, e12)) | ~ spl42_40), inference(avatar_component_clause, [], [f710])).
fof(f710, plain, (spl42_40 <=> (e13 = op1(e11, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_40])])).
fof(f3107, plain, (~ (op2(e21, e22) = h4(op1(e11, e12))) | (~ spl42_160 | spl42_202)), inference(forward_demodulation, [], [f1672, f1445])).
fof(f1672, plain, (~ (h4(op1(e11, e12)) = op2(e21, h4(e12))) | spl42_202), inference(avatar_component_clause, [], [f1670])).
fof(f1670, plain, (spl42_202 <=> (h4(op1(e11, e12)) = op2(e21, h4(e12)))), introduced(avatar_definition, [new_symbols(naming, [spl42_202])])).
fof(f3078, plain, (~ spl42_35 | spl42_201), inference(avatar_contradiction_clause, [], [f3077])).
fof(f3077, plain, ($false | (~ spl42_35 | spl42_201)), inference(trivial_inequality_removal, [], [f3076])).
fof(f3076, plain, (~ (h4(e12) = h4(e12)) | (~ spl42_35 | spl42_201)), inference(forward_demodulation, [], [f1668, f691])).
fof(f691, plain, ((e12 = op1(e11, e13)) | ~ spl42_35), inference(avatar_component_clause, [], [f689])).
fof(f689, plain, (spl42_35 <=> (e12 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_35])])).
fof(f1668, plain, (~ (h4(e12) = h4(op1(e11, e13))) | spl42_201), inference(avatar_component_clause, [], [f1666])).
fof(f1666, plain, (spl42_201 <=> (h4(e12) = h4(op1(e11, e13)))), introduced(avatar_definition, [new_symbols(naming, [spl42_201])])).
fof(f3054, plain, (~ spl42_31 | ~ spl42_95 | ~ spl42_160 | ~ spl42_167 | spl42_193), inference(avatar_contradiction_clause, [], [f3053])).
fof(f3053, plain, ($false | (~ spl42_31 | ~ spl42_95 | ~ spl42_160 | ~ spl42_167 | spl42_193)), inference(subsumption_resolution, [], [f3052, f978])).
fof(f978, plain, ((e22 = op2(e22, e20)) | ~ spl42_95), inference(avatar_component_clause, [], [f976])).
fof(f976, plain, (spl42_95 <=> (e22 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_95])])).
fof(f3052, plain, (~ (e22 = op2(e22, e20)) | (~ spl42_31 | ~ spl42_160 | ~ spl42_167 | spl42_193)), inference(forward_demodulation, [], [f3051, f1445])).
fof(f3051, plain, (~ (op2(e22, e20) = h4(e12)) | (~ spl42_31 | ~ spl42_160 | ~ spl42_167 | spl42_193)), inference(forward_demodulation, [], [f3050, f674])).
fof(f674, plain, ((e12 = op1(e12, e10)) | ~ spl42_31), inference(avatar_component_clause, [], [f672])).
fof(f672, plain, (spl42_31 <=> (e12 = op1(e12, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_31])])).
fof(f3050, plain, (~ (op2(e22, e20) = h4(op1(e12, e10))) | (~ spl42_160 | ~ spl42_167 | spl42_193)), inference(forward_demodulation, [], [f3049, f1445])).
fof(f3049, plain, (~ (h4(op1(e12, e10)) = op2(h4(e12), e20)) | (~ spl42_167 | spl42_193)), inference(forward_demodulation, [], [f1636, f1484])).
fof(f1636, plain, (~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | spl42_193), inference(avatar_component_clause, [], [f1634])).
fof(f1634, plain, (spl42_193 <=> (h4(op1(e12, e10)) = op2(h4(e12), h4(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl42_193])])).
fof(f3021, plain, (~ spl42_16 | ~ spl42_80 | ~ spl42_167 | spl42_198), inference(avatar_contradiction_clause, [], [f3020])).
fof(f3020, plain, ($false | (~ spl42_16 | ~ spl42_80 | ~ spl42_167 | spl42_198)), inference(subsumption_resolution, [], [f3019, f914])).
fof(f914, plain, ((e23 = op2(e23, e20)) | ~ spl42_80), inference(avatar_component_clause, [], [f912])).
fof(f912, plain, (spl42_80 <=> (e23 = op2(e23, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_80])])).
fof(f3019, plain, (~ (e23 = op2(e23, e20)) | (~ spl42_16 | ~ spl42_167 | spl42_198)), inference(forward_demodulation, [], [f3018, f476])).
fof(f3018, plain, (~ (op2(e23, e20) = h4(e13)) | (~ spl42_16 | ~ spl42_167 | spl42_198)), inference(forward_demodulation, [], [f3017, f610])).
fof(f610, plain, ((e13 = op1(e13, e10)) | ~ spl42_16), inference(avatar_component_clause, [], [f608])).
fof(f608, plain, (spl42_16 <=> (e13 = op1(e13, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_16])])).
fof(f3017, plain, (~ (op2(e23, e20) = h4(op1(e13, e10))) | (~ spl42_167 | spl42_198)), inference(forward_demodulation, [], [f1656, f1484])).
fof(f1656, plain, (~ (h4(op1(e13, e10)) = op2(e23, h4(e10))) | spl42_198), inference(avatar_component_clause, [], [f1654])).
fof(f1654, plain, (spl42_198 <=> (h4(op1(e13, e10)) = op2(e23, h4(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl42_198])])).
fof(f2980, plain, (~ spl42_21 | ~ spl42_85 | ~ spl42_160 | ~ spl42_167 | spl42_194), inference(avatar_contradiction_clause, [], [f2979])).
fof(f2979, plain, ($false | (~ spl42_21 | ~ spl42_85 | ~ spl42_160 | ~ spl42_167 | spl42_194)), inference(subsumption_resolution, [], [f2978, f936])).
fof(f936, plain, ((e20 = op2(e22, e22)) | ~ spl42_85), inference(avatar_component_clause, [], [f934])).
fof(f934, plain, (spl42_85 <=> (e20 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_85])])).
fof(f2978, plain, (~ (e20 = op2(e22, e22)) | (~ spl42_21 | ~ spl42_160 | ~ spl42_167 | spl42_194)), inference(forward_demodulation, [], [f2977, f1484])).
fof(f2977, plain, (~ (op2(e22, e22) = h4(e10)) | (~ spl42_21 | ~ spl42_160 | spl42_194)), inference(forward_demodulation, [], [f2976, f632])).
fof(f632, plain, ((e10 = op1(e12, e12)) | ~ spl42_21), inference(avatar_component_clause, [], [f630])).
fof(f630, plain, (spl42_21 <=> (e10 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_21])])).
fof(f2976, plain, (~ (op2(e22, e22) = h4(op1(e12, e12))) | (~ spl42_160 | spl42_194)), inference(forward_demodulation, [], [f1640, f1445])).
fof(f1640, plain, (~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | spl42_194), inference(avatar_component_clause, [], [f1638])).
fof(f1638, plain, (spl42_194 <=> (h4(op1(e12, e12)) = op2(h4(e12), h4(e12)))), introduced(avatar_definition, [new_symbols(naming, [spl42_194])])).
fof(f2933, plain, (~ spl42_55 | ~ spl42_119 | ~ spl42_160 | ~ spl42_167 | spl42_192), inference(avatar_contradiction_clause, [], [f2932])).
fof(f2932, plain, ($false | (~ spl42_55 | ~ spl42_119 | ~ spl42_160 | ~ spl42_167 | spl42_192)), inference(subsumption_resolution, [], [f2931, f1080])).
fof(f1080, plain, ((e22 = op2(e20, e22)) | ~ spl42_119), inference(avatar_component_clause, [], [f1078])).
fof(f1078, plain, (spl42_119 <=> (e22 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_119])])).
fof(f2931, plain, (~ (e22 = op2(e20, e22)) | (~ spl42_55 | ~ spl42_160 | ~ spl42_167 | spl42_192)), inference(forward_demodulation, [], [f2930, f1445])).
fof(f2930, plain, (~ (op2(e20, e22) = h4(e12)) | (~ spl42_55 | ~ spl42_160 | ~ spl42_167 | spl42_192)), inference(forward_demodulation, [], [f2929, f776])).
fof(f776, plain, ((e12 = op1(e10, e12)) | ~ spl42_55), inference(avatar_component_clause, [], [f774])).
fof(f774, plain, (spl42_55 <=> (e12 = op1(e10, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_55])])).
fof(f2929, plain, (~ (op2(e20, e22) = h4(op1(e10, e12))) | (~ spl42_160 | ~ spl42_167 | spl42_192)), inference(forward_demodulation, [], [f2928, f1484])).
fof(f2928, plain, (~ (h4(op1(e10, e12)) = op2(h4(e10), e22)) | (~ spl42_160 | spl42_192)), inference(forward_demodulation, [], [f1632, f1445])).
fof(f1632, plain, (~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | spl42_192), inference(avatar_component_clause, [], [f1630])).
fof(f1630, plain, (spl42_192 <=> (h4(op1(e10, e12)) = op2(h4(e10), h4(e12)))), introduced(avatar_definition, [new_symbols(naming, [spl42_192])])).
fof(f2899, plain, (spl42_190 | ~ spl42_106 | ~ spl42_191), inference(avatar_split_clause, [], [f2897, f1605, f1023, f1600])).
fof(f1023, plain, (spl42_106 <=> (e21 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_106])])).
fof(f2897, plain, ((e21 = h1(e11)) | (~ spl42_106 | ~ spl42_191)), inference(backward_demodulation, [], [f2827, f1025])).
fof(f1025, plain, ((e21 = op2(e21, e21)) | ~ spl42_106), inference(avatar_component_clause, [], [f1023])).
fof(f2892, plain, (~ spl42_71 | ~ spl42_119), inference(avatar_split_clause, [], [f2889, f1078, f874])).
fof(f874, plain, (spl42_71 <=> (e22 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_71])])).
fof(f2889, plain, (~ (e22 = op2(e23, e22)) | ~ spl42_119), inference(backward_demodulation, [], [f266, f1080])).
fof(f266, plain, ~ (op2(e20, e22) = op2(e23, e22)), inference(cnf_transformation, [], [f6])).
fof(f6, plain, (~ (op2(e23, e22) = op2(e23, e23)) & ~ (op2(e23, e21) = op2(e23, e23)) & ~ (op2(e23, e21) = op2(e23, e22)) & ~ (op2(e23, e20) = op2(e23, e23)) & ~ (op2(e23, e20) = op2(e23, e22)) & ~ (op2(e23, e20) = op2(e23, e21)) & ~ (op2(e22, e22) = op2(e22, e23)) & ~ (op2(e22, e21) = op2(e22, e23)) & ~ (op2(e22, e21) = op2(e22, e22)) & ~ (op2(e22, e20) = op2(e22, e23)) & ~ (op2(e22, e20) = op2(e22, e22)) & ~ (op2(e22, e20) = op2(e22, e21)) & ~ (op2(e21, e22) = op2(e21, e23)) & ~ (op2(e21, e21) = op2(e21, e23)) & ~ (op2(e21, e21) = op2(e21, e22)) & ~ (op2(e21, e20) = op2(e21, e23)) & ~ (op2(e21, e20) = op2(e21, e22)) & ~ (op2(e21, e20) = op2(e21, e21)) & ~ (op2(e20, e22) = op2(e20, e23)) & ~ (op2(e20, e21) = op2(e20, e23)) & ~ (op2(e20, e21) = op2(e20, e22)) & ~ (op2(e20, e20) = op2(e20, e23)) & ~ (op2(e20, e20) = op2(e20, e22)) & ~ (op2(e20, e20) = op2(e20, e21)) & ~ (op2(e22, e23) = op2(e23, e23)) & ~ (op2(e21, e23) = op2(e23, e23)) & ~ (op2(e21, e23) = op2(e22, e23)) & ~ (op2(e20, e23) = op2(e23, e23)) & ~ (op2(e20, e23) = op2(e22, e23)) & ~ (op2(e20, e23) = op2(e21, e23)) & ~ (op2(e22, e22) = op2(e23, e22)) & ~ (op2(e21, e22) = op2(e23, e22)) & ~ (op2(e21, e22) = op2(e22, e22)) & ~ (op2(e20, e22) = op2(e23, e22)) & ~ (op2(e20, e22) = op2(e22, e22)) & ~ (op2(e20, e22) = op2(e21, e22)) & ~ (op2(e22, e21) = op2(e23, e21)) & ~ (op2(e21, e21) = op2(e23, e21)) & ~ (op2(e21, e21) = op2(e22, e21)) & ~ (op2(e20, e21) = op2(e23, e21)) & ~ (op2(e20, e21) = op2(e22, e21)) & ~ (op2(e20, e21) = op2(e21, e21)) & ~ (op2(e22, e20) = op2(e23, e20)) & ~ (op2(e21, e20) = op2(e23, e20)) & ~ (op2(e21, e20) = op2(e22, e20)) & ~ (op2(e20, e20) = op2(e23, e20)) & ~ (op2(e20, e20) = op2(e22, e20)) & ~ (op2(e20, e20) = op2(e21, e20))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG121+1.p', ax6)).
fof(f2881, plain, (~ spl42_18 | ~ spl42_82 | ~ spl42_160 | spl42_199), inference(avatar_contradiction_clause, [], [f2880])).
fof(f2880, plain, ($false | (~ spl42_18 | ~ spl42_82 | ~ spl42_160 | spl42_199)), inference(subsumption_resolution, [], [f2879, f923])).
fof(f923, plain, ((e21 = op2(e22, e23)) | ~ spl42_82), inference(avatar_component_clause, [], [f921])).
fof(f921, plain, (spl42_82 <=> (e21 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_82])])).
fof(f2879, plain, (~ (e21 = op2(e22, e23)) | (~ spl42_18 | ~ spl42_160 | spl42_199)), inference(forward_demodulation, [], [f2878, f1430])).
fof(f2878, plain, (~ (op2(e22, e23) = h4(e11)) | (~ spl42_18 | ~ spl42_160 | spl42_199)), inference(forward_demodulation, [], [f2877, f619])).
fof(f619, plain, ((e11 = op1(e12, e13)) | ~ spl42_18), inference(avatar_component_clause, [], [f617])).
fof(f617, plain, (spl42_18 <=> (e11 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_18])])).
fof(f2877, plain, (~ (op2(e22, e23) = h4(op1(e12, e13))) | (~ spl42_160 | spl42_199)), inference(forward_demodulation, [], [f1660, f1445])).
fof(f1660, plain, (~ (h4(op1(e12, e13)) = op2(h4(e12), e23)) | spl42_199), inference(avatar_component_clause, [], [f1658])).
fof(f1658, plain, (spl42_199 <=> (h4(op1(e12, e13)) = op2(h4(e12), e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_199])])).
fof(f2873, plain, (~ spl42_118 | ~ spl42_191), inference(avatar_split_clause, [], [f2872, f1605, f1074])).
fof(f1074, plain, (spl42_118 <=> (e21 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_118])])).
fof(f2872, plain, (~ (e21 = op2(e20, e22)) | ~ spl42_191), inference(forward_demodulation, [], [f1405, f1606])).
fof(f1405, plain, ~ (op2(e20, e22) = h1(e10)), inference(backward_demodulation, [], [f277, f465])).
fof(f277, plain, ~ (op2(e20, e20) = op2(e20, e22)), inference(cnf_transformation, [], [f6])).
fof(f2823, plain, (spl42_191 | ~ spl42_126), inference(avatar_split_clause, [], [f2650, f1108, f1605])).
fof(f1108, plain, (spl42_126 <=> (op2(e20, e20) = e21)), introduced(avatar_definition, [new_symbols(naming, [spl42_126])])).
fof(f2650, plain, ((e21 = h1(e10)) | ~ spl42_126), inference(backward_demodulation, [], [f465, f1110])).
fof(f1110, plain, ((op2(e20, e20) = e21) | ~ spl42_126), inference(avatar_component_clause, [], [f1108])).
fof(f2807, plain, (~ spl42_115 | ~ spl42_160), inference(avatar_split_clause, [], [f2806, f1444, f1061])).
fof(f1061, plain, (spl42_115 <=> (e22 = op2(e20, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_115])])).
fof(f2806, plain, (~ (e22 = op2(e20, e23)) | ~ spl42_160), inference(forward_demodulation, [], [f1425, f1445])).
fof(f1425, plain, ~ (op2(e20, e23) = h4(e12)), inference(backward_demodulation, [], [f270, f1424])).
fof(f1424, plain, (op2(e21, e23) = h4(e12)), inference(forward_demodulation, [], [f479, f462])).
fof(f479, plain, (op2(op2(op2(e23, e23), op2(e23, e23)), e23) = h4(e12)), inference(cnf_transformation, [], [f17])).
fof(f270, plain, ~ (op2(e20, e23) = op2(e21, e23)), inference(cnf_transformation, [], [f6])).
fof(f2803, plain, (~ spl42_109 | ~ spl42_45 | ~ spl42_167 | spl42_204), inference(avatar_split_clause, [], [f2802, f1678, f1483, f732, f1036])).
fof(f1036, plain, (spl42_109 <=> (e20 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_109])])).
fof(f732, plain, (spl42_45 <=> (e10 = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_45])])).
fof(f1678, plain, (spl42_204 <=> (h4(op1(e11, e10)) = op2(e21, h4(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl42_204])])).
fof(f2802, plain, (~ (e20 = op2(e21, e20)) | (~ spl42_45 | ~ spl42_167 | spl42_204)), inference(forward_demodulation, [], [f2801, f1484])).
fof(f2801, plain, (~ (op2(e21, e20) = h4(e10)) | (~ spl42_45 | ~ spl42_167 | spl42_204)), inference(forward_demodulation, [], [f2463, f734])).
fof(f734, plain, ((e10 = op1(e11, e10)) | ~ spl42_45), inference(avatar_component_clause, [], [f732])).
fof(f2463, plain, (~ (op2(e21, e20) = h4(op1(e11, e10))) | (~ spl42_167 | spl42_204)), inference(backward_demodulation, [], [f1680, f1484])).
fof(f1680, plain, (~ (h4(op1(e11, e10)) = op2(e21, h4(e10))) | spl42_204), inference(avatar_component_clause, [], [f1678])).
fof(f2798, plain, (~ spl42_111 | ~ spl42_160), inference(avatar_split_clause, [], [f2797, f1444, f1044])).
fof(f1044, plain, (spl42_111 <=> (e22 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_111])])).
fof(f2797, plain, (~ (e22 = op2(e21, e20)) | ~ spl42_160), inference(forward_demodulation, [], [f1428, f1445])).
fof(f1428, plain, ~ (op2(e21, e20) = h4(e12)), inference(backward_demodulation, [], [f284, f1424])).
fof(f284, plain, ~ (op2(e21, e20) = op2(e21, e23)), inference(cnf_transformation, [], [f6])).
fof(f2795, plain, (~ spl42_80 | ~ spl42_112), inference(avatar_contradiction_clause, [], [f2794])).
fof(f2794, plain, ($false | (~ spl42_80 | ~ spl42_112)), inference(subsumption_resolution, [], [f2793, f1050])).
fof(f1050, plain, ((e23 = op2(e21, e20)) | ~ spl42_112), inference(avatar_component_clause, [], [f1048])).
fof(f1048, plain, (spl42_112 <=> (e23 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_112])])).
fof(f2793, plain, (~ (e23 = op2(e21, e20)) | ~ spl42_80), inference(forward_demodulation, [], [f256, f914])).
fof(f256, plain, ~ (op2(e21, e20) = op2(e23, e20)), inference(cnf_transformation, [], [f6])).
fof(f2773, plain, (~ spl42_95 | ~ spl42_96), inference(avatar_contradiction_clause, [], [f2772])).
fof(f2772, plain, ($false | (~ spl42_95 | ~ spl42_96)), inference(subsumption_resolution, [], [f2771, f311])).
fof(f311, plain, ~ (e22 = e23), inference(cnf_transformation, [], [f8])).
fof(f8, plain, (~ (e22 = e23) & ~ (e21 = e23) & ~ (e21 = e22) & ~ (e20 = e23) & ~ (e20 = e22) & ~ (e20 = e21)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG121+1.p', ax8)).
fof(f2771, plain, ((e22 = e23) | (~ spl42_95 | ~ spl42_96)), inference(forward_demodulation, [], [f982, f978])).
fof(f982, plain, ((e23 = op2(e22, e20)) | ~ spl42_96), inference(avatar_component_clause, [], [f980])).
fof(f980, plain, (spl42_96 <=> (e23 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_96])])).
fof(f2768, plain, (~ spl42_90 | ~ spl42_82), inference(avatar_split_clause, [], [f2767, f921, f955])).
fof(f955, plain, (spl42_90 <=> (e21 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_90])])).
fof(f2767, plain, (~ (e21 = op2(e22, e21)) | ~ spl42_82), inference(forward_demodulation, [], [f292, f923])).
fof(f292, plain, ~ (op2(e22, e21) = op2(e22, e23)), inference(cnf_transformation, [], [f6])).
fof(f2764, plain, (~ spl42_81 | ~ spl42_167), inference(avatar_split_clause, [], [f2616, f1483, f917])).
fof(f917, plain, (spl42_81 <=> (e20 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_81])])).
fof(f2616, plain, (~ (e20 = op2(e22, e23)) | ~ spl42_167), inference(forward_demodulation, [], [f1432, f1484])).
fof(f1432, plain, ~ (op2(e22, e23) = h4(e10)), inference(backward_demodulation, [], [f275, f477])).
fof(f477, plain, (op2(e23, e23) = h4(e10)), inference(cnf_transformation, [], [f17])).
fof(f275, plain, ~ (op2(e22, e23) = op2(e23, e23)), inference(cnf_transformation, [], [f6])).
fof(f2763, plain, (~ spl42_69 | ~ spl42_167), inference(avatar_split_clause, [], [f2603, f1483, f866])).
fof(f866, plain, (spl42_69 <=> (e20 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_69])])).
fof(f2603, plain, (~ (e20 = op2(e23, e22)) | ~ spl42_167), inference(forward_demodulation, [], [f1435, f1484])).
fof(f1435, plain, ~ (op2(e23, e22) = h4(e10)), inference(backward_demodulation, [], [f299, f477])).
fof(f299, plain, ~ (op2(e23, e22) = op2(e23, e23)), inference(cnf_transformation, [], [f6])).
fof(f2711, plain, (~ spl42_57 | ~ spl42_60), inference(avatar_contradiction_clause, [], [f2710])).
fof(f2710, plain, ($false | (~ spl42_57 | ~ spl42_60)), inference(subsumption_resolution, [], [f2709, f302])).
fof(f302, plain, ~ (e10 = e13), inference(cnf_transformation, [], [f7])).
fof(f2709, plain, ((e10 = e13) | (~ spl42_57 | ~ spl42_60)), inference(backward_demodulation, [], [f797, f785])).
fof(f797, plain, ((e13 = op1(e10, e11)) | ~ spl42_60), inference(avatar_component_clause, [], [f795])).
fof(f795, plain, (spl42_60 <=> (e13 = op1(e10, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_60])])).
fof(f2706, plain, (~ spl42_54 | ~ spl42_62), inference(avatar_split_clause, [], [f2702, f804, f770])).
fof(f770, plain, (spl42_54 <=> (e11 = op1(e10, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_54])])).
fof(f2702, plain, (~ (e11 = op1(e10, e12)) | ~ spl42_62), inference(backward_demodulation, [], [f229, f806])).
fof(f229, plain, ~ (op1(e10, e10) = op1(e10, e12)), inference(cnf_transformation, [], [f5])).
fof(f5, plain, (~ (op1(e13, e12) = op1(e13, e13)) & ~ (op1(e13, e11) = op1(e13, e13)) & ~ (op1(e13, e11) = op1(e13, e12)) & ~ (op1(e13, e10) = op1(e13, e13)) & ~ (op1(e13, e10) = op1(e13, e12)) & ~ (op1(e13, e10) = op1(e13, e11)) & ~ (op1(e12, e12) = op1(e12, e13)) & ~ (op1(e12, e11) = op1(e12, e13)) & ~ (op1(e12, e11) = op1(e12, e12)) & ~ (op1(e12, e10) = op1(e12, e13)) & ~ (op1(e12, e10) = op1(e12, e12)) & ~ (op1(e12, e10) = op1(e12, e11)) & ~ (op1(e11, e12) = op1(e11, e13)) & ~ (op1(e11, e11) = op1(e11, e13)) & ~ (op1(e11, e11) = op1(e11, e12)) & ~ (op1(e11, e10) = op1(e11, e13)) & ~ (op1(e11, e10) = op1(e11, e12)) & ~ (op1(e11, e10) = op1(e11, e11)) & ~ (op1(e10, e12) = op1(e10, e13)) & ~ (op1(e10, e11) = op1(e10, e13)) & ~ (op1(e10, e11) = op1(e10, e12)) & ~ (op1(e10, e10) = op1(e10, e13)) & ~ (op1(e10, e10) = op1(e10, e12)) & ~ (op1(e10, e10) = op1(e10, e11)) & ~ (op1(e12, e13) = op1(e13, e13)) & ~ (op1(e11, e13) = op1(e13, e13)) & ~ (op1(e11, e13) = op1(e12, e13)) & ~ (op1(e10, e13) = op1(e13, e13)) & ~ (op1(e10, e13) = op1(e12, e13)) & ~ (op1(e10, e13) = op1(e11, e13)) & ~ (op1(e12, e12) = op1(e13, e12)) & ~ (op1(e11, e12) = op1(e13, e12)) & ~ (op1(e11, e12) = op1(e12, e12)) & ~ (op1(e10, e12) = op1(e13, e12)) & ~ (op1(e10, e12) = op1(e12, e12)) & ~ (op1(e10, e12) = op1(e11, e12)) & ~ (op1(e12, e11) = op1(e13, e11)) & ~ (op1(e11, e11) = op1(e13, e11)) & ~ (op1(e11, e11) = op1(e12, e11)) & ~ (op1(e10, e11) = op1(e13, e11)) & ~ (op1(e10, e11) = op1(e12, e11)) & ~ (op1(e10, e11) = op1(e11, e11)) & ~ (op1(e12, e10) = op1(e13, e10)) & ~ (op1(e11, e10) = op1(e13, e10)) & ~ (op1(e11, e10) = op1(e12, e10)) & ~ (op1(e10, e10) = op1(e13, e10)) & ~ (op1(e10, e10) = op1(e12, e10)) & ~ (op1(e10, e10) = op1(e11, e10))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG121+1.p', ax5)).
fof(f2705, plain, (~ spl42_46 | ~ spl42_62), inference(avatar_split_clause, [], [f2701, f804, f736])).
fof(f736, plain, (spl42_46 <=> (e11 = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_46])])).
fof(f2701, plain, (~ (e11 = op1(e11, e10)) | ~ spl42_62), inference(backward_demodulation, [], [f204, f806])).
fof(f204, plain, ~ (op1(e10, e10) = op1(e11, e10)), inference(cnf_transformation, [], [f5])).
fof(f2697, plain, (~ spl42_70 | ~ spl42_74), inference(avatar_split_clause, [], [f2696, f887, f870])).
fof(f887, plain, (spl42_74 <=> (e21 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_74])])).
fof(f2696, plain, (~ (e21 = op2(e23, e22)) | ~ spl42_74), inference(backward_demodulation, [], [f297, f889])).
fof(f889, plain, ((e21 = op2(e23, e21)) | ~ spl42_74), inference(avatar_component_clause, [], [f887])).
fof(f297, plain, ~ (op2(e23, e21) = op2(e23, e22)), inference(cnf_transformation, [], [f6])).
fof(f2679, plain, (~ spl42_93 | ~ spl42_95), inference(avatar_contradiction_clause, [], [f2678])).
fof(f2678, plain, ($false | (~ spl42_93 | ~ spl42_95)), inference(subsumption_resolution, [], [f2677, f307])).
fof(f307, plain, ~ (e20 = e22), inference(cnf_transformation, [], [f8])).
fof(f2677, plain, ((e20 = e22) | (~ spl42_93 | ~ spl42_95)), inference(forward_demodulation, [], [f978, f970])).
fof(f970, plain, ((e20 = op2(e22, e20)) | ~ spl42_93), inference(avatar_component_clause, [], [f968])).
fof(f968, plain, (spl42_93 <=> (e20 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_93])])).
fof(f2647, plain, (~ spl42_83 | ~ spl42_160), inference(avatar_split_clause, [], [f2645, f1444, f925])).
fof(f925, plain, (spl42_83 <=> (e22 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_83])])).
fof(f2645, plain, (~ (e22 = op2(e22, e23)) | ~ spl42_160), inference(backward_demodulation, [], [f1426, f1445])).
fof(f1426, plain, ~ (op2(e22, e23) = h4(e12)), inference(backward_demodulation, [], [f273, f1424])).
fof(f273, plain, ~ (op2(e21, e23) = op2(e22, e23)), inference(cnf_transformation, [], [f6])).
fof(f2640, plain, (~ spl42_120 | ~ spl42_116), inference(avatar_split_clause, [], [f2639, f1065, f1082])).
fof(f1082, plain, (spl42_120 <=> (e23 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_120])])).
fof(f2639, plain, (~ (e23 = op2(e20, e22)) | ~ spl42_116), inference(forward_demodulation, [], [f281, f1067])).
fof(f281, plain, ~ (op2(e20, e22) = op2(e20, e23)), inference(cnf_transformation, [], [f6])).
fof(f2584, plain, (~ spl42_39 | ~ spl42_35), inference(avatar_split_clause, [], [f2583, f689, f706])).
fof(f706, plain, (spl42_39 <=> (e12 = op1(e11, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_39])])).
fof(f2583, plain, (~ (e12 = op1(e11, e12)) | ~ spl42_35), inference(forward_demodulation, [], [f239, f691])).
fof(f239, plain, ~ (op1(e11, e12) = op1(e11, e13)), inference(cnf_transformation, [], [f5])).
fof(f2582, plain, (~ spl42_38 | ~ spl42_42), inference(avatar_split_clause, [], [f2581, f719, f702])).
fof(f702, plain, (spl42_38 <=> (e11 = op1(e11, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_38])])).
fof(f2581, plain, (~ (e11 = op1(e11, e12)) | ~ spl42_42), inference(forward_demodulation, [], [f237, f721])).
fof(f237, plain, ~ (op1(e11, e11) = op1(e11, e12)), inference(cnf_transformation, [], [f5])).
fof(f2578, plain, (~ spl42_27 | ~ spl42_31), inference(avatar_contradiction_clause, [], [f2577])).
fof(f2577, plain, ($false | (~ spl42_27 | ~ spl42_31)), inference(subsumption_resolution, [], [f2576, f674])).
fof(f2576, plain, (~ (e12 = op1(e12, e10)) | ~ spl42_27), inference(forward_demodulation, [], [f240, f657])).
fof(f657, plain, ((e12 = op1(e12, e11)) | ~ spl42_27), inference(avatar_component_clause, [], [f655])).
fof(f655, plain, (spl42_27 <=> (e12 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_27])])).
fof(f240, plain, ~ (op1(e12, e10) = op1(e12, e11)), inference(cnf_transformation, [], [f5])).
fof(f2565, plain, (~ spl42_6 | ~ spl42_7), inference(avatar_contradiction_clause, [], [f2564])).
fof(f2564, plain, ($false | (~ spl42_6 | ~ spl42_7)), inference(subsumption_resolution, [], [f2563, f303])).
fof(f303, plain, ~ (e11 = e12), inference(cnf_transformation, [], [f7])).
fof(f2563, plain, ((e11 = e12) | (~ spl42_6 | ~ spl42_7)), inference(forward_demodulation, [], [f572, f568])).
fof(f572, plain, ((e12 = op1(e13, e12)) | ~ spl42_7), inference(avatar_component_clause, [], [f570])).
fof(f570, plain, (spl42_7 <=> (e12 = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_7])])).
fof(f2540, plain, (~ spl42_23 | ~ spl42_31), inference(avatar_split_clause, [], [f2538, f672, f638])).
fof(f638, plain, (spl42_23 <=> (e12 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_23])])).
fof(f2538, plain, (~ (e12 = op1(e12, e12)) | ~ spl42_31), inference(backward_demodulation, [], [f241, f674])).
fof(f241, plain, ~ (op1(e12, e10) = op1(e12, e12)), inference(cnf_transformation, [], [f5])).
fof(f2537, plain, (~ spl42_35 | ~ spl42_36), inference(avatar_contradiction_clause, [], [f2536])).
fof(f2536, plain, ($false | (~ spl42_35 | ~ spl42_36)), inference(subsumption_resolution, [], [f2535, f305])).
fof(f305, plain, ~ (e12 = e13), inference(cnf_transformation, [], [f7])).
fof(f2535, plain, ((e12 = e13) | (~ spl42_35 | ~ spl42_36)), inference(forward_demodulation, [], [f695, f691])).
fof(f695, plain, ((e13 = op1(e11, e13)) | ~ spl42_36), inference(avatar_component_clause, [], [f693])).
fof(f693, plain, (spl42_36 <=> (e13 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_36])])).
fof(f2479, plain, (~ spl42_191 | ~ spl42_110), inference(avatar_split_clause, [], [f2476, f1040, f1605])).
fof(f1040, plain, (spl42_110 <=> (e21 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_110])])).
fof(f2476, plain, (~ (e21 = h1(e10)) | ~ spl42_110), inference(backward_demodulation, [], [f1401, f1042])).
fof(f1042, plain, ((e21 = op2(e21, e20)) | ~ spl42_110), inference(avatar_component_clause, [], [f1040])).
fof(f1401, plain, ~ (op2(e21, e20) = h1(e10)), inference(backward_demodulation, [], [f252, f465])).
fof(f252, plain, ~ (op2(e20, e20) = op2(e21, e20)), inference(cnf_transformation, [], [f6])).
fof(f2472, plain, (~ spl42_191 | ~ spl42_122), inference(avatar_split_clause, [], [f2470, f1091, f1605])).
fof(f1091, plain, (spl42_122 <=> (e21 = op2(e20, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_122])])).
fof(f2470, plain, (~ (e21 = h1(e10)) | ~ spl42_122), inference(backward_demodulation, [], [f1404, f1093])).
fof(f1093, plain, ((e21 = op2(e20, e21)) | ~ spl42_122), inference(avatar_component_clause, [], [f1091])).
fof(f1404, plain, ~ (op2(e20, e21) = h1(e10)), inference(backward_demodulation, [], [f276, f465])).
fof(f276, plain, ~ (op2(e20, e20) = op2(e20, e21)), inference(cnf_transformation, [], [f6])).
fof(f2465, plain, (spl42_126 | ~ spl42_167), inference(avatar_split_clause, [], [f2461, f1483, f1108])).
fof(f2461, plain, ((op2(e20, e20) = e21) | ~ spl42_167), inference(backward_demodulation, [], [f1436, f1484])).
fof(f1436, plain, (e21 = op2(h4(e10), h4(e10))), inference(backward_demodulation, [], [f462, f477])).
fof(f2456, plain, (~ spl42_191 | ~ spl42_114), inference(avatar_split_clause, [], [f2455, f1057, f1605])).
fof(f1057, plain, (spl42_114 <=> (e21 = op2(e20, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_114])])).
fof(f2455, plain, (~ (e21 = h1(e10)) | ~ spl42_114), inference(forward_demodulation, [], [f1406, f1059])).
fof(f1059, plain, ((e21 = op2(e20, e23)) | ~ spl42_114), inference(avatar_component_clause, [], [f1057])).
fof(f1406, plain, ~ (op2(e20, e23) = h1(e10)), inference(backward_demodulation, [], [f278, f465])).
fof(f278, plain, ~ (op2(e20, e20) = op2(e20, e23)), inference(cnf_transformation, [], [f6])).
fof(f2447, plain, (~ spl42_121 | ~ spl42_117), inference(avatar_split_clause, [], [f2446, f1070, f1087])).
fof(f1070, plain, (spl42_117 <=> (e20 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_117])])).
fof(f2446, plain, (~ (e20 = op2(e20, e21)) | ~ spl42_117), inference(forward_demodulation, [], [f279, f1072])).
fof(f1072, plain, ((e20 = op2(e20, e22)) | ~ spl42_117), inference(avatar_component_clause, [], [f1070])).
fof(f279, plain, ~ (op2(e20, e21) = op2(e20, e22)), inference(cnf_transformation, [], [f6])).
fof(f2432, plain, (spl42_160 | ~ spl42_99), inference(avatar_split_clause, [], [f2431, f993, f1444])).
fof(f993, plain, (spl42_99 <=> (e22 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_99])])).
fof(f2431, plain, ((e22 = h4(e12)) | ~ spl42_99), inference(forward_demodulation, [], [f1424, f995])).
fof(f995, plain, ((e22 = op2(e21, e23)) | ~ spl42_99), inference(avatar_component_clause, [], [f993])).
fof(f2405, plain, (spl42_62 | ~ spl42_1), inference(avatar_split_clause, [], [f2351, f545, f804])).
fof(f545, plain, (spl42_1 <=> (e10 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_1])])).
fof(f2351, plain, ((op1(e10, e10) = e11) | ~ spl42_1), inference(backward_demodulation, [], [f459, f547])).
fof(f547, plain, ((e10 = op1(e13, e13)) | ~ spl42_1), inference(avatar_component_clause, [], [f545])).
fof(f459, plain, (e11 = op1(op1(e13, e13), op1(e13, e13))), inference(cnf_transformation, [], [f12])).
fof(f12, plain, ((e12 = op1(op1(op1(e13, e13), op1(e13, e13)), e13)) & (e11 = op1(op1(e13, e13), op1(e13, e13))) & (e10 = op1(e13, e13))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG121+1.p', ax12)).
fof(f2391, plain, (~ spl42_49 | ~ spl42_1), inference(avatar_split_clause, [], [f2390, f545, f749])).
fof(f749, plain, (spl42_49 <=> (e10 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_49])])).
fof(f2390, plain, (~ (e10 = op1(e10, e13)) | ~ spl42_1), inference(forward_demodulation, [], [f224, f547])).
fof(f224, plain, ~ (op1(e10, e13) = op1(e13, e13)), inference(cnf_transformation, [], [f5])).
fof(f2389, plain, (~ spl42_48 | ~ spl42_16), inference(avatar_split_clause, [], [f2388, f608, f744])).
fof(f744, plain, (spl42_48 <=> (e13 = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_48])])).
fof(f2388, plain, (~ (e13 = op1(e11, e10)) | ~ spl42_16), inference(forward_demodulation, [], [f208, f610])).
fof(f208, plain, ~ (op1(e11, e10) = op1(e13, e10)), inference(cnf_transformation, [], [f5])).
fof(f2385, plain, (~ spl42_47 | ~ spl42_35), inference(avatar_split_clause, [], [f2384, f689, f740])).
fof(f740, plain, (spl42_47 <=> (e12 = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_47])])).
fof(f2384, plain, (~ (e12 = op1(e11, e10)) | ~ spl42_35), inference(forward_demodulation, [], [f236, f691])).
fof(f236, plain, ~ (op1(e11, e10) = op1(e11, e13)), inference(cnf_transformation, [], [f5])).
fof(f2383, plain, (~ spl42_42 | ~ spl42_26), inference(avatar_split_clause, [], [f2382, f651, f719])).
fof(f651, plain, (spl42_26 <=> (e11 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_26])])).
fof(f2382, plain, (~ (e11 = op1(e11, e11)) | ~ spl42_26), inference(forward_demodulation, [], [f213, f653])).
fof(f653, plain, ((e11 = op1(e12, e11)) | ~ spl42_26), inference(avatar_component_clause, [], [f651])).
fof(f213, plain, ~ (op1(e11, e11) = op1(e12, e11)), inference(cnf_transformation, [], [f5])).
fof(f2364, plain, (~ spl42_17 | ~ spl42_1), inference(avatar_split_clause, [], [f2363, f545, f613])).
fof(f613, plain, (spl42_17 <=> (e10 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_17])])).
fof(f2363, plain, (~ (e10 = op1(e12, e13)) | ~ spl42_1), inference(forward_demodulation, [], [f227, f547])).
fof(f227, plain, ~ (op1(e12, e13) = op1(e13, e13)), inference(cnf_transformation, [], [f5])).
fof(f2357, plain, (~ spl42_1 | spl42_195), inference(avatar_contradiction_clause, [], [f2356])).
fof(f2356, plain, ($false | (~ spl42_1 | spl42_195)), inference(trivial_inequality_removal, [], [f2355])).
fof(f2355, plain, (~ (h4(e10) = h4(e10)) | (~ spl42_1 | spl42_195)), inference(backward_demodulation, [], [f1644, f547])).
fof(f1644, plain, (~ (h4(e10) = h4(op1(e13, e13))) | spl42_195), inference(avatar_component_clause, [], [f1642])).
fof(f1642, plain, (spl42_195 <=> (h4(e10) = h4(op1(e13, e13)))), introduced(avatar_definition, [new_symbols(naming, [spl42_195])])).
fof(f2341, plain, (~ spl42_15 | ~ spl42_16), inference(avatar_contradiction_clause, [], [f2340])).
fof(f2340, plain, ($false | (~ spl42_15 | ~ spl42_16)), inference(subsumption_resolution, [], [f2339, f305])).
fof(f2339, plain, ((e12 = e13) | (~ spl42_15 | ~ spl42_16)), inference(backward_demodulation, [], [f610, f606])).
fof(f606, plain, ((e12 = op1(e13, e10)) | ~ spl42_15), inference(avatar_component_clause, [], [f604])).
fof(f604, plain, (spl42_15 <=> (e12 = op1(e13, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_15])])).
fof(f2338, plain, (~ spl42_4 | ~ spl42_16), inference(avatar_split_clause, [], [f2335, f608, f557])).
fof(f557, plain, (spl42_4 <=> (e13 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_4])])).
fof(f2335, plain, (~ (e13 = op1(e13, e13)) | ~ spl42_16), inference(backward_demodulation, [], [f248, f610])).
fof(f248, plain, ~ (op1(e13, e10) = op1(e13, e13)), inference(cnf_transformation, [], [f5])).
fof(f2336, plain, (~ spl42_12 | ~ spl42_16), inference(avatar_split_clause, [], [f2333, f608, f591])).
fof(f591, plain, (spl42_12 <=> (e13 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_12])])).
fof(f2333, plain, (~ (e13 = op1(e13, e11)) | ~ spl42_16), inference(backward_demodulation, [], [f246, f610])).
fof(f246, plain, ~ (op1(e13, e10) = op1(e13, e11)), inference(cnf_transformation, [], [f5])).
fof(f2305, plain, (~ spl42_34 | ~ spl42_35), inference(avatar_contradiction_clause, [], [f2304])).
fof(f2304, plain, ($false | (~ spl42_34 | ~ spl42_35)), inference(subsumption_resolution, [], [f2303, f303])).
fof(f2303, plain, ((e11 = e12) | (~ spl42_34 | ~ spl42_35)), inference(backward_demodulation, [], [f691, f687])).
fof(f687, plain, ((e11 = op1(e11, e13)) | ~ spl42_34), inference(avatar_component_clause, [], [f685])).
fof(f685, plain, (spl42_34 <=> (e11 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_34])])).
fof(f2300, plain, (~ spl42_3 | ~ spl42_35), inference(avatar_split_clause, [], [f2297, f689, f553])).
fof(f553, plain, (spl42_3 <=> (e12 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_3])])).
fof(f2297, plain, (~ (e12 = op1(e13, e13)) | ~ spl42_35), inference(backward_demodulation, [], [f226, f691])).
fof(f226, plain, ~ (op1(e11, e13) = op1(e13, e13)), inference(cnf_transformation, [], [f5])).
fof(f2277, plain, (~ spl42_29 | ~ spl42_45), inference(avatar_split_clause, [], [f2272, f732, f664])).
fof(f664, plain, (spl42_29 <=> (e10 = op1(e12, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_29])])).
fof(f2272, plain, (~ (e10 = op1(e12, e10)) | ~ spl42_45), inference(backward_demodulation, [], [f207, f734])).
fof(f207, plain, ~ (op1(e11, e10) = op1(e12, e10)), inference(cnf_transformation, [], [f5])).
fof(f2259, plain, (~ spl42_40 | ~ spl42_56), inference(avatar_split_clause, [], [f2255, f778, f710])).
fof(f778, plain, (spl42_56 <=> (e13 = op1(e10, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_56])])).
fof(f2255, plain, (~ (e13 = op1(e11, e12)) | ~ spl42_56), inference(backward_demodulation, [], [f216, f780])).
fof(f780, plain, ((e13 = op1(e10, e12)) | ~ spl42_56), inference(avatar_component_clause, [], [f778])).
fof(f216, plain, ~ (op1(e10, e12) = op1(e11, e12)), inference(cnf_transformation, [], [f5])).
fof(f2253, plain, (~ spl42_53 | ~ spl42_57), inference(avatar_split_clause, [], [f2247, f783, f766])).
fof(f766, plain, (spl42_53 <=> (e10 = op1(e10, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_53])])).
fof(f2247, plain, (~ (e10 = op1(e10, e12)) | ~ spl42_57), inference(backward_demodulation, [], [f231, f785])).
fof(f231, plain, ~ (op1(e10, e11) = op1(e10, e12)), inference(cnf_transformation, [], [f5])).
fof(f2251, plain, (~ spl42_25 | ~ spl42_57), inference(avatar_split_clause, [], [f2245, f783, f647])).
fof(f647, plain, (spl42_25 <=> (e10 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_25])])).
fof(f2245, plain, (~ (e10 = op1(e12, e11)) | ~ spl42_57), inference(backward_demodulation, [], [f211, f785])).
fof(f211, plain, ~ (op1(e10, e11) = op1(e12, e11)), inference(cnf_transformation, [], [f5])).
fof(f2239, plain, (~ spl42_30 | ~ spl42_62), inference(avatar_split_clause, [], [f2231, f804, f668])).
fof(f668, plain, (spl42_30 <=> (e11 = op1(e12, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_30])])).
fof(f2231, plain, (~ (e11 = op1(e12, e10)) | ~ spl42_62), inference(backward_demodulation, [], [f205, f806])).
fof(f205, plain, ~ (op1(e10, e10) = op1(e12, e10)), inference(cnf_transformation, [], [f5])).
fof(f2229, plain, (spl42_167 | ~ spl42_65), inference(avatar_split_clause, [], [f2228, f849, f1483])).
fof(f849, plain, (spl42_65 <=> (e20 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_65])])).
fof(f2228, plain, ((e20 = h4(e10)) | ~ spl42_65), inference(backward_demodulation, [], [f477, f851])).
fof(f851, plain, ((e20 = op2(e23, e23)) | ~ spl42_65), inference(avatar_component_clause, [], [f849])).
fof(f2206, plain, (~ spl42_79 | ~ spl42_80), inference(avatar_contradiction_clause, [], [f2205])).
fof(f2205, plain, ($false | (~ spl42_79 | ~ spl42_80)), inference(subsumption_resolution, [], [f2204, f311])).
fof(f2204, plain, ((e22 = e23) | (~ spl42_79 | ~ spl42_80)), inference(backward_demodulation, [], [f914, f910])).
fof(f910, plain, ((e22 = op2(e23, e20)) | ~ spl42_79), inference(avatar_component_clause, [], [f908])).
fof(f908, plain, (spl42_79 <=> (e22 = op2(e23, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_79])])).
fof(f2203, plain, (~ spl42_72 | ~ spl42_80), inference(avatar_split_clause, [], [f2200, f912, f878])).
fof(f878, plain, (spl42_72 <=> (e23 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_72])])).
fof(f2200, plain, (~ (e23 = op2(e23, e22)) | ~ spl42_80), inference(backward_demodulation, [], [f295, f914])).
fof(f295, plain, ~ (op2(e23, e20) = op2(e23, e22)), inference(cnf_transformation, [], [f6])).
fof(f2125, plain, (~ spl42_84 | ~ spl42_116), inference(avatar_split_clause, [], [f2122, f1065, f929])).
fof(f929, plain, (spl42_84 <=> (e23 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_84])])).
fof(f2122, plain, (~ (e23 = op2(e22, e23)) | ~ spl42_116), inference(backward_demodulation, [], [f271, f1067])).
fof(f271, plain, ~ (op2(e20, e23) = op2(e22, e23)), inference(cnf_transformation, [], [f6])).
fof(f2111, plain, (~ spl42_113 | ~ spl42_121), inference(avatar_split_clause, [], [f2106, f1087, f1053])).
fof(f1053, plain, (spl42_113 <=> (e20 = op2(e20, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_113])])).
fof(f2106, plain, (~ (e20 = op2(e20, e23)) | ~ spl42_121), inference(backward_demodulation, [], [f280, f1089])).
fof(f280, plain, ~ (op2(e20, e21) = op2(e20, e23)), inference(cnf_transformation, [], [f6])).
fof(f2108, plain, (~ spl42_89 | ~ spl42_121), inference(avatar_split_clause, [], [f2103, f1087, f951])).
fof(f951, plain, (spl42_89 <=> (e20 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_89])])).
fof(f2103, plain, (~ (e20 = op2(e22, e21)) | ~ spl42_121), inference(backward_demodulation, [], [f259, f1089])).
fof(f259, plain, ~ (op2(e20, e21) = op2(e22, e21)), inference(cnf_transformation, [], [f6])).
fof(f1693, plain, (~ spl42_192 | ~ spl42_193 | ~ spl42_194 | spl42_165 | spl42_162 | spl42_159 | ~ spl42_195 | ~ spl42_196 | ~ spl42_197 | ~ spl42_198 | ~ spl42_199 | ~ spl42_200 | ~ spl42_201 | ~ spl42_202 | ~ spl42_203 | ~ spl42_204 | ~ spl42_205 | ~ spl42_206 | ~ spl42_207), inference(avatar_split_clause, [], [f1628, f1690, f1686, f1682, f1678, f1674, f1670, f1666, f1662, f1658, f1654, f1650, f1646, f1642, f1440, f1456, f1473, f1638, f1634, f1630])).
fof(f1473, plain, (spl42_165 <=> sP39), introduced(avatar_definition, [new_symbols(naming, [spl42_165])])).
fof(f1456, plain, (spl42_162 <=> sP40), introduced(avatar_definition, [new_symbols(naming, [spl42_162])])).
fof(f1440, plain, (spl42_159 <=> sP41), introduced(avatar_definition, [new_symbols(naming, [spl42_159])])).
fof(f1628, plain, (~ (e21 = h4(op1(e10, e10))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), e21)) | ~ (h4(op1(e10, e13)) = op2(h4(e10), e23)) | ~ (h4(op1(e11, e10)) = op2(e21, h4(e10))) | ~ (h2(e10) = h4(op1(e11, e11))) | ~ (h4(op1(e11, e12)) = op2(e21, h4(e12))) | ~ (h4(e12) = h4(op1(e11, e13))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), e21)) | ~ (h4(op1(e12, e13)) = op2(h4(e12), e23)) | ~ (h4(op1(e13, e10)) = op2(e23, h4(e10))) | ~ (op2(e23, e21) = h4(op1(e13, e11))) | ~ (h4(op1(e13, e12)) = op2(e23, h4(e12))) | ~ (h4(e10) = h4(op1(e13, e13))) | sP41 | sP40 | sP39 | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12)))), inference(forward_demodulation, [], [f1627, f1436])).
fof(f1627, plain, (~ (h4(op1(e10, e11)) = op2(h4(e10), e21)) | ~ (h4(op1(e10, e13)) = op2(h4(e10), e23)) | ~ (h4(op1(e11, e10)) = op2(e21, h4(e10))) | ~ (h2(e10) = h4(op1(e11, e11))) | ~ (h4(op1(e11, e12)) = op2(e21, h4(e12))) | ~ (h4(e12) = h4(op1(e11, e13))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), e21)) | ~ (h4(op1(e12, e13)) = op2(h4(e12), e23)) | ~ (h4(op1(e13, e10)) = op2(e23, h4(e10))) | ~ (op2(e23, e21) = h4(op1(e13, e11))) | ~ (h4(op1(e13, e12)) = op2(e23, h4(e12))) | ~ (h4(e10) = h4(op1(e13, e13))) | sP41 | sP40 | sP39 | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f1626, f1430])).
fof(f1626, plain, (~ (h4(op1(e10, e13)) = op2(h4(e10), e23)) | ~ (h4(op1(e11, e10)) = op2(e21, h4(e10))) | ~ (h2(e10) = h4(op1(e11, e11))) | ~ (h4(op1(e11, e12)) = op2(e21, h4(e12))) | ~ (h4(e12) = h4(op1(e11, e13))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), e21)) | ~ (h4(op1(e12, e13)) = op2(h4(e12), e23)) | ~ (h4(op1(e13, e10)) = op2(e23, h4(e10))) | ~ (op2(e23, e21) = h4(op1(e13, e11))) | ~ (h4(op1(e13, e12)) = op2(e23, h4(e12))) | ~ (h4(e10) = h4(op1(e13, e13))) | sP41 | sP40 | sP39 | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f1625, f476])).
fof(f1625, plain, (~ (h4(op1(e11, e10)) = op2(e21, h4(e10))) | ~ (h2(e10) = h4(op1(e11, e11))) | ~ (h4(op1(e11, e12)) = op2(e21, h4(e12))) | ~ (h4(e12) = h4(op1(e11, e13))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), e21)) | ~ (h4(op1(e12, e13)) = op2(h4(e12), e23)) | ~ (h4(op1(e13, e10)) = op2(e23, h4(e10))) | ~ (op2(e23, e21) = h4(op1(e13, e11))) | ~ (h4(op1(e13, e12)) = op2(e23, h4(e12))) | ~ (h4(e10) = h4(op1(e13, e13))) | sP41 | sP40 | sP39 | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f1624, f1430])).
fof(f1624, plain, (~ (h2(e10) = h4(op1(e11, e11))) | ~ (h4(op1(e11, e12)) = op2(e21, h4(e12))) | ~ (h4(e12) = h4(op1(e11, e13))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), e21)) | ~ (h4(op1(e12, e13)) = op2(h4(e12), e23)) | ~ (h4(op1(e13, e10)) = op2(e23, h4(e10))) | ~ (op2(e23, e21) = h4(op1(e13, e11))) | ~ (h4(op1(e13, e12)) = op2(e23, h4(e12))) | ~ (h4(e10) = h4(op1(e13, e13))) | sP41 | sP40 | sP39 | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f1623, f469])).
fof(f1623, plain, (~ (op2(e21, e21) = h4(op1(e11, e11))) | ~ (h4(op1(e11, e12)) = op2(e21, h4(e12))) | ~ (h4(e12) = h4(op1(e11, e13))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), e21)) | ~ (h4(op1(e12, e13)) = op2(h4(e12), e23)) | ~ (h4(op1(e13, e10)) = op2(e23, h4(e10))) | ~ (op2(e23, e21) = h4(op1(e13, e11))) | ~ (h4(op1(e13, e12)) = op2(e23, h4(e12))) | ~ (h4(e10) = h4(op1(e13, e13))) | sP41 | sP40 | sP39 | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f1622, f1430])).
fof(f1622, plain, (~ (h4(op1(e11, e12)) = op2(e21, h4(e12))) | ~ (h4(e12) = h4(op1(e11, e13))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), e21)) | ~ (h4(op1(e12, e13)) = op2(h4(e12), e23)) | ~ (h4(op1(e13, e10)) = op2(e23, h4(e10))) | ~ (op2(e23, e21) = h4(op1(e13, e11))) | ~ (h4(op1(e13, e12)) = op2(e23, h4(e12))) | ~ (h4(e10) = h4(op1(e13, e13))) | sP41 | sP40 | sP39 | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f1621, f1430])).
fof(f1621, plain, (~ (h4(e12) = h4(op1(e11, e13))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), e21)) | ~ (h4(op1(e12, e13)) = op2(h4(e12), e23)) | ~ (h4(op1(e13, e10)) = op2(e23, h4(e10))) | ~ (op2(e23, e21) = h4(op1(e13, e11))) | ~ (h4(op1(e13, e12)) = op2(e23, h4(e12))) | ~ (h4(e10) = h4(op1(e13, e13))) | sP41 | sP40 | sP39 | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f1620, f1424])).
fof(f1620, plain, (~ (op2(e21, e23) = h4(op1(e11, e13))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), e21)) | ~ (h4(op1(e12, e13)) = op2(h4(e12), e23)) | ~ (h4(op1(e13, e10)) = op2(e23, h4(e10))) | ~ (op2(e23, e21) = h4(op1(e13, e11))) | ~ (h4(op1(e13, e12)) = op2(e23, h4(e12))) | ~ (h4(e10) = h4(op1(e13, e13))) | sP41 | sP40 | sP39 | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f1619, f1430])).
fof(f1619, plain, (~ (h4(op1(e11, e13)) = op2(h4(e11), e23)) | ~ (h4(op1(e12, e11)) = op2(h4(e12), e21)) | ~ (h4(op1(e12, e13)) = op2(h4(e12), e23)) | ~ (h4(op1(e13, e10)) = op2(e23, h4(e10))) | ~ (op2(e23, e21) = h4(op1(e13, e11))) | ~ (h4(op1(e13, e12)) = op2(e23, h4(e12))) | ~ (h4(e10) = h4(op1(e13, e13))) | sP41 | sP40 | sP39 | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f1618, f476])).
fof(f1618, plain, (~ (h4(op1(e12, e11)) = op2(h4(e12), e21)) | ~ (h4(op1(e12, e13)) = op2(h4(e12), e23)) | ~ (h4(op1(e13, e10)) = op2(e23, h4(e10))) | ~ (op2(e23, e21) = h4(op1(e13, e11))) | ~ (h4(op1(e13, e12)) = op2(e23, h4(e12))) | ~ (h4(e10) = h4(op1(e13, e13))) | sP41 | sP40 | sP39 | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f1617, f1430])).
fof(f1617, plain, (~ (h4(op1(e12, e13)) = op2(h4(e12), e23)) | ~ (h4(op1(e13, e10)) = op2(e23, h4(e10))) | ~ (op2(e23, e21) = h4(op1(e13, e11))) | ~ (h4(op1(e13, e12)) = op2(e23, h4(e12))) | ~ (h4(e10) = h4(op1(e13, e13))) | sP41 | sP40 | sP39 | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f1616, f476])).
fof(f1616, plain, (~ (h4(op1(e13, e10)) = op2(e23, h4(e10))) | ~ (op2(e23, e21) = h4(op1(e13, e11))) | ~ (h4(op1(e13, e12)) = op2(e23, h4(e12))) | ~ (h4(e10) = h4(op1(e13, e13))) | sP41 | sP40 | sP39 | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f1615, f476])).
fof(f1615, plain, (~ (op2(e23, e21) = h4(op1(e13, e11))) | ~ (h4(op1(e13, e12)) = op2(e23, h4(e12))) | ~ (h4(e10) = h4(op1(e13, e13))) | sP41 | sP40 | sP39 | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f1614, f476])).
fof(f1614, plain, (~ (h4(op1(e13, e11)) = op2(h4(e13), e21)) | ~ (h4(op1(e13, e12)) = op2(e23, h4(e12))) | ~ (h4(e10) = h4(op1(e13, e13))) | sP41 | sP40 | sP39 | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f1613, f1430])).
fof(f1613, plain, (~ (h4(op1(e13, e12)) = op2(e23, h4(e12))) | ~ (h4(e10) = h4(op1(e13, e13))) | sP41 | sP40 | sP39 | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f1612, f476])).
fof(f1612, plain, (~ (h4(e10) = h4(op1(e13, e13))) | sP41 | sP40 | sP39 | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f1611, f477])).
fof(f1611, plain, (~ (op2(e23, e23) = h4(op1(e13, e13))) | sP41 | sP40 | sP39 | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f1610, f476])).
fof(f1610, plain, (sP41 | sP40 | sP39 | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(subsumption_resolution, [], [f543, f476])).
fof(f543, plain, (~ (e23 = h4(e13)) | sP41 | sP40 | sP39 | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(cnf_transformation, [], [f65])).
fof(f65, plain, (((~ (e23 = h4(e13)) & ~ (e23 = h4(e12)) & ~ (e23 = h4(e11)) & ~ (e23 = h4(e10))) | sP41 | sP40 | sP39 | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) & ((~ (e23 = h3(e13)) & ~ (e23 = h3(e12)) & ~ (e23 = h3(e11)) & ~ (e23 = h3(e10))) | sP38 | sP37 | sP36 | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) & ((~ (e23 = h2(e13)) & ~ (e23 = h2(e12)) & ~ (e23 = h2(e11)) & ~ (e23 = h2(e10))) | sP35 | sP34 | sP33 | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) & ((~ (e23 = h1(e13)) & ~ (e23 = h1(e12)) & ~ (e23 = h1(e11)) & ~ (e23 = h1(e10))) | sP32 | sP31 | sP30 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(definition_folding, [], [f20, e64, e63, e62, e61, e60, e59, e58, e57, e56, e55, e54, e53])).
fof(f53, plain, ((~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10))) | ~ sP30), inference(usedef, [], [e53])).
fof(e53, plain, (sP30 <=> (~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP30])])).
fof(f54, plain, ((~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10))) | ~ sP31), inference(usedef, [], [e54])).
fof(e54, plain, (sP31 <=> (~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP31])])).
fof(f55, plain, ((~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10))) | ~ sP32), inference(usedef, [], [e55])).
fof(e55, plain, (sP32 <=> (~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP32])])).
fof(f56, plain, ((~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10))) | ~ sP33), inference(usedef, [], [e56])).
fof(e56, plain, (sP33 <=> (~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP33])])).
fof(f57, plain, ((~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10))) | ~ sP34), inference(usedef, [], [e57])).
fof(e57, plain, (sP34 <=> (~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP34])])).
fof(f58, plain, ((~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10))) | ~ sP35), inference(usedef, [], [e58])).
fof(e58, plain, (sP35 <=> (~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP35])])).
fof(f59, plain, ((~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10))) | ~ sP36), inference(usedef, [], [e59])).
fof(e59, plain, (sP36 <=> (~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP36])])).
fof(f60, plain, ((~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10))) | ~ sP37), inference(usedef, [], [e60])).
fof(e60, plain, (sP37 <=> (~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP37])])).
fof(f61, plain, ((~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10))) | ~ sP38), inference(usedef, [], [e61])).
fof(e61, plain, (sP38 <=> (~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP38])])).
fof(f62, plain, ((~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10))) | ~ sP39), inference(usedef, [], [e62])).
fof(e62, plain, (sP39 <=> (~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP39])])).
fof(f63, plain, ((~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10))) | ~ sP40), inference(usedef, [], [e63])).
fof(e63, plain, (sP40 <=> (~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP40])])).
fof(f64, plain, ((~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10))) | ~ sP41), inference(usedef, [], [e64])).
fof(e64, plain, (sP41 <=> (~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP41])])).
fof(f20, plain, (((~ (e23 = h4(e13)) & ~ (e23 = h4(e12)) & ~ (e23 = h4(e11)) & ~ (e23 = h4(e10))) | (~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10))) | (~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10))) | (~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10))) | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) & ((~ (e23 = h3(e13)) & ~ (e23 = h3(e12)) & ~ (e23 = h3(e11)) & ~ (e23 = h3(e10))) | (~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10))) | (~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10))) | (~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10))) | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) & ((~ (e23 = h2(e13)) & ~ (e23 = h2(e12)) & ~ (e23 = h2(e11)) & ~ (e23 = h2(e10))) | (~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10))) | (~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10))) | (~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) & ((~ (e23 = h1(e13)) & ~ (e23 = h1(e12)) & ~ (e23 = h1(e11)) & ~ (e23 = h1(e10))) | (~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10))) | (~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10))) | (~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10))) | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(ennf_transformation, [], [f19])).
fof(f19, plain, ~ ((((e23 = h4(e13)) | (e23 = h4(e12)) | (e23 = h4(e11)) | (e23 = h4(e10))) & ((e22 = h4(e13)) | (e22 = h4(e12)) | (e22 = h4(e11)) | (e22 = h4(e10))) & ((e21 = h4(e13)) | (e21 = h4(e12)) | (e21 = h4(e11)) | (e21 = h4(e10))) & ((e20 = h4(e13)) | (e20 = h4(e12)) | (e20 = h4(e11)) | (e20 = h4(e10))) & (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) & (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) & (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) & (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) & (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) & (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) & (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) & (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) & (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) & (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) & (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) & (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) & (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) & (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) & (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) & (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) | (((e23 = h3(e13)) | (e23 = h3(e12)) | (e23 = h3(e11)) | (e23 = h3(e10))) & ((e22 = h3(e13)) | (e22 = h3(e12)) | (e22 = h3(e11)) | (e22 = h3(e10))) & ((e21 = h3(e13)) | (e21 = h3(e12)) | (e21 = h3(e11)) | (e21 = h3(e10))) & ((e20 = h3(e13)) | (e20 = h3(e12)) | (e20 = h3(e11)) | (e20 = h3(e10))) & (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) & (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) & (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) & (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) & (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) & (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) & (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) & (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) & (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) & (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) & (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) & (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) & (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) & (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) & (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) & (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) | (((e23 = h2(e13)) | (e23 = h2(e12)) | (e23 = h2(e11)) | (e23 = h2(e10))) & ((e22 = h2(e13)) | (e22 = h2(e12)) | (e22 = h2(e11)) | (e22 = h2(e10))) & ((e21 = h2(e13)) | (e21 = h2(e12)) | (e21 = h2(e11)) | (e21 = h2(e10))) & ((e20 = h2(e13)) | (e20 = h2(e12)) | (e20 = h2(e11)) | (e20 = h2(e10))) & (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) & (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) & (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) & (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) & (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) & (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) & (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) & (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) & (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) & (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) & (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) & (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) & (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) & (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) & (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) & (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) | (((e23 = h1(e13)) | (e23 = h1(e12)) | (e23 = h1(e11)) | (e23 = h1(e10))) & ((e22 = h1(e13)) | (e22 = h1(e12)) | (e22 = h1(e11)) | (e22 = h1(e10))) & ((e21 = h1(e13)) | (e21 = h1(e12)) | (e21 = h1(e11)) | (e21 = h1(e10))) & ((e20 = h1(e13)) | (e20 = h1(e12)) | (e20 = h1(e11)) | (e20 = h1(e10))) & (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) & (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) & (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) & (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) & (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) & (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) & (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) & (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) & (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) & (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) & (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) & (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) & (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) & (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) & (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) & (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(negated_conjecture, [], [f18])).
fof(f18, plain, ~ ((((e23 = h4(e13)) | (e23 = h4(e12)) | (e23 = h4(e11)) | (e23 = h4(e10))) & ((e22 = h4(e13)) | (e22 = h4(e12)) | (e22 = h4(e11)) | (e22 = h4(e10))) & ((e21 = h4(e13)) | (e21 = h4(e12)) | (e21 = h4(e11)) | (e21 = h4(e10))) & ((e20 = h4(e13)) | (e20 = h4(e12)) | (e20 = h4(e11)) | (e20 = h4(e10))) & (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) & (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) & (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) & (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) & (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) & (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) & (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) & (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) & (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) & (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) & (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) & (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) & (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) & (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) & (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) & (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) | (((e23 = h3(e13)) | (e23 = h3(e12)) | (e23 = h3(e11)) | (e23 = h3(e10))) & ((e22 = h3(e13)) | (e22 = h3(e12)) | (e22 = h3(e11)) | (e22 = h3(e10))) & ((e21 = h3(e13)) | (e21 = h3(e12)) | (e21 = h3(e11)) | (e21 = h3(e10))) & ((e20 = h3(e13)) | (e20 = h3(e12)) | (e20 = h3(e11)) | (e20 = h3(e10))) & (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) & (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) & (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) & (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) & (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) & (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) & (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) & (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) & (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) & (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) & (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) & (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) & (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) & (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) & (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) & (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) | (((e23 = h2(e13)) | (e23 = h2(e12)) | (e23 = h2(e11)) | (e23 = h2(e10))) & ((e22 = h2(e13)) | (e22 = h2(e12)) | (e22 = h2(e11)) | (e22 = h2(e10))) & ((e21 = h2(e13)) | (e21 = h2(e12)) | (e21 = h2(e11)) | (e21 = h2(e10))) & ((e20 = h2(e13)) | (e20 = h2(e12)) | (e20 = h2(e11)) | (e20 = h2(e10))) & (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) & (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) & (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) & (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) & (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) & (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) & (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) & (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) & (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) & (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) & (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) & (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) & (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) & (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) & (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) & (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) | (((e23 = h1(e13)) | (e23 = h1(e12)) | (e23 = h1(e11)) | (e23 = h1(e10))) & ((e22 = h1(e13)) | (e22 = h1(e12)) | (e22 = h1(e11)) | (e22 = h1(e10))) & ((e21 = h1(e13)) | (e21 = h1(e12)) | (e21 = h1(e11)) | (e21 = h1(e10))) & ((e20 = h1(e13)) | (e20 = h1(e12)) | (e20 = h1(e11)) | (e20 = h1(e10))) & (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) & (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) & (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) & (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) & (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) & (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) & (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) & (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) & (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) & (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) & (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) & (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) & (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) & (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) & (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) & (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG121+1.p', co1)).
fof(f1486, plain, (~ spl42_165 | ~ spl42_167), inference(avatar_split_clause, [], [f488, f1483, f1473])).
fof(f488, plain, (~ (e20 = h4(e10)) | ~ sP39), inference(cnf_transformation, [], [f98])).
fof(f98, plain, ((~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10))) | ~ sP39), inference(nnf_transformation, [], [f62])).
fof(f1465, plain, ~ spl42_162, inference(avatar_split_clause, [], [f1464, f1456])).
fof(f1464, plain, ~ sP40, inference(subsumption_resolution, [], [f485, f1430])).
fof(f485, plain, (~ (e21 = h4(e11)) | ~ sP40), inference(cnf_transformation, [], [f97])).
fof(f97, plain, ((~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10))) | ~ sP40), inference(nnf_transformation, [], [f63])).
fof(f1447, plain, (~ spl42_159 | ~ spl42_160), inference(avatar_split_clause, [], [f482, f1444, f1440])).
fof(f482, plain, (~ (e22 = h4(e12)) | ~ sP41), inference(cnf_transformation, [], [f96])).
fof(f96, plain, ((~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10))) | ~ sP41), inference(nnf_transformation, [], [f64])).
fof(f1399, plain, spl42_65, inference(avatar_split_clause, [], [f461, f849])).
fof(f461, plain, (e20 = op2(e23, e23)), inference(cnf_transformation, [], [f13])).
fof(f1398, plain, spl42_99, inference(avatar_split_clause, [], [f1397, f993])).
fof(f1397, plain, (e22 = op2(e21, e23)), inference(backward_demodulation, [], [f463, f462])).
fof(f463, plain, (e22 = op2(op2(op2(e23, e23), op2(e23, e23)), e23)), inference(cnf_transformation, [], [f13])).
fof(f1396, plain, spl42_1, inference(avatar_split_clause, [], [f458, f545])).
fof(f458, plain, (e10 = op1(e13, e13)), inference(cnf_transformation, [], [f12])).
fof(f1395, plain, spl42_35, inference(avatar_split_clause, [], [f1394, f689])).
fof(f1394, plain, (e12 = op1(e11, e13)), inference(backward_demodulation, [], [f460, f459])).
fof(f460, plain, (e12 = op1(op1(op1(e13, e13), op1(e13, e13)), e13)), inference(cnf_transformation, [], [f12])).
fof(f1392, plain, (~ spl42_126 | spl42_121), inference(avatar_split_clause, [], [f440, f1087, f1108])).
fof(f440, plain, ((e20 = op2(e20, e21)) | ~ (op2(e20, e20) = e21)), inference(cnf_transformation, [], [f52])).
fof(f52, plain, (((~ (e23 = op2(e23, e23)) & (e23 = op2(e23, e23)) & (e23 = op2(e23, e23))) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15) & ((e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23))) & ((e23 = op2(e23, e22)) | ~ (e22 = op2(e23, e23))) & ((e23 = op2(e23, e21)) | ~ (e21 = op2(e23, e23))) & ((e23 = op2(e23, e20)) | ~ (e20 = op2(e23, e23))) & ((e22 = op2(e22, e23)) | ~ (e23 = op2(e22, e22))) & ((e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & ((e22 = op2(e22, e21)) | ~ (e21 = op2(e22, e22))) & ((e22 = op2(e22, e20)) | ~ (e20 = op2(e22, e22))) & ((e21 = op2(e21, e23)) | ~ (e23 = op2(e21, e21))) & ((e21 = op2(e21, e22)) | ~ (e22 = op2(e21, e21))) & ((e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & ((e21 = op2(e21, e20)) | ~ (e20 = op2(e21, e21))) & ((e20 = op2(e20, e23)) | ~ (op2(e20, e20) = e23)) & ((e20 = op2(e20, e22)) | ~ (op2(e20, e20) = e22)) & ((e20 = op2(e20, e21)) | ~ (op2(e20, e20) = e21)) & ((e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20))) & ((e23 = op2(e23, e23)) | (e22 = op2(e22, e22)) | (e21 = op2(e21, e21)) | (e20 = op2(e20, e20)))), inference(definition_folding, [], [f11, e51, e50, e49, e48, e47, e46, e45, e44, e43, e42, e41, e40, e39, e38, e37])).
fof(f37, plain, ((~ (e20 = op2(e20, e20)) & (e20 = op2(e20, e20)) & (e20 = op2(e20, e20))) | ~ sP15), inference(usedef, [], [e37])).
fof(e37, plain, (sP15 <=> (~ (e20 = op2(e20, e20)) & (e20 = op2(e20, e20)) & (e20 = op2(e20, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f38, plain, ((~ (e21 = op2(e21, e21)) & (e20 = op2(e21, e20)) & (e20 = op2(e20, e21))) | ~ sP16), inference(usedef, [], [e38])).
fof(e38, plain, (sP16 <=> (~ (e21 = op2(e21, e21)) & (e20 = op2(e21, e20)) & (e20 = op2(e20, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f39, plain, ((~ (e22 = op2(e22, e22)) & (e20 = op2(e22, e20)) & (e20 = op2(e20, e22))) | ~ sP17), inference(usedef, [], [e39])).
fof(e39, plain, (sP17 <=> (~ (e22 = op2(e22, e22)) & (e20 = op2(e22, e20)) & (e20 = op2(e20, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f40, plain, ((~ (e23 = op2(e23, e23)) & (e20 = op2(e23, e20)) & (e20 = op2(e20, e23))) | ~ sP18), inference(usedef, [], [e40])).
fof(e40, plain, (sP18 <=> (~ (e23 = op2(e23, e23)) & (e20 = op2(e23, e20)) & (e20 = op2(e20, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f41, plain, ((~ (e20 = op2(e20, e20)) & (e21 = op2(e20, e21)) & (e21 = op2(e21, e20))) | ~ sP19), inference(usedef, [], [e41])).
fof(e41, plain, (sP19 <=> (~ (e20 = op2(e20, e20)) & (e21 = op2(e20, e21)) & (e21 = op2(e21, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP19])])).
fof(f42, plain, ((~ (e21 = op2(e21, e21)) & (e21 = op2(e21, e21)) & (e21 = op2(e21, e21))) | ~ sP20), inference(usedef, [], [e42])).
fof(e42, plain, (sP20 <=> (~ (e21 = op2(e21, e21)) & (e21 = op2(e21, e21)) & (e21 = op2(e21, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP20])])).
fof(f43, plain, ((~ (e22 = op2(e22, e22)) & (e21 = op2(e22, e21)) & (e21 = op2(e21, e22))) | ~ sP21), inference(usedef, [], [e43])).
fof(e43, plain, (sP21 <=> (~ (e22 = op2(e22, e22)) & (e21 = op2(e22, e21)) & (e21 = op2(e21, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP21])])).
fof(f44, plain, ((~ (e23 = op2(e23, e23)) & (e21 = op2(e23, e21)) & (e21 = op2(e21, e23))) | ~ sP22), inference(usedef, [], [e44])).
fof(e44, plain, (sP22 <=> (~ (e23 = op2(e23, e23)) & (e21 = op2(e23, e21)) & (e21 = op2(e21, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP22])])).
fof(f45, plain, ((~ (e20 = op2(e20, e20)) & (e22 = op2(e20, e22)) & (e22 = op2(e22, e20))) | ~ sP23), inference(usedef, [], [e45])).
fof(e45, plain, (sP23 <=> (~ (e20 = op2(e20, e20)) & (e22 = op2(e20, e22)) & (e22 = op2(e22, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP23])])).
fof(f46, plain, ((~ (e21 = op2(e21, e21)) & (e22 = op2(e21, e22)) & (e22 = op2(e22, e21))) | ~ sP24), inference(usedef, [], [e46])).
fof(e46, plain, (sP24 <=> (~ (e21 = op2(e21, e21)) & (e22 = op2(e21, e22)) & (e22 = op2(e22, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP24])])).
fof(f47, plain, ((~ (e22 = op2(e22, e22)) & (e22 = op2(e22, e22)) & (e22 = op2(e22, e22))) | ~ sP25), inference(usedef, [], [e47])).
fof(e47, plain, (sP25 <=> (~ (e22 = op2(e22, e22)) & (e22 = op2(e22, e22)) & (e22 = op2(e22, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP25])])).
fof(f48, plain, ((~ (e23 = op2(e23, e23)) & (e22 = op2(e23, e22)) & (e22 = op2(e22, e23))) | ~ sP26), inference(usedef, [], [e48])).
fof(e48, plain, (sP26 <=> (~ (e23 = op2(e23, e23)) & (e22 = op2(e23, e22)) & (e22 = op2(e22, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP26])])).
fof(f49, plain, ((~ (e20 = op2(e20, e20)) & (e23 = op2(e20, e23)) & (e23 = op2(e23, e20))) | ~ sP27), inference(usedef, [], [e49])).
fof(e49, plain, (sP27 <=> (~ (e20 = op2(e20, e20)) & (e23 = op2(e20, e23)) & (e23 = op2(e23, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP27])])).
fof(f50, plain, ((~ (e21 = op2(e21, e21)) & (e23 = op2(e21, e23)) & (e23 = op2(e23, e21))) | ~ sP28), inference(usedef, [], [e50])).
fof(e50, plain, (sP28 <=> (~ (e21 = op2(e21, e21)) & (e23 = op2(e21, e23)) & (e23 = op2(e23, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP28])])).
fof(f51, plain, ((~ (e22 = op2(e22, e22)) & (e23 = op2(e22, e23)) & (e23 = op2(e23, e22))) | ~ sP29), inference(usedef, [], [e51])).
fof(e51, plain, (sP29 <=> (~ (e22 = op2(e22, e22)) & (e23 = op2(e22, e23)) & (e23 = op2(e23, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP29])])).
fof(f11, plain, (((~ (e23 = op2(e23, e23)) & (e23 = op2(e23, e23)) & (e23 = op2(e23, e23))) | (~ (e22 = op2(e22, e22)) & (e23 = op2(e22, e23)) & (e23 = op2(e23, e22))) | (~ (e21 = op2(e21, e21)) & (e23 = op2(e21, e23)) & (e23 = op2(e23, e21))) | (~ (e20 = op2(e20, e20)) & (e23 = op2(e20, e23)) & (e23 = op2(e23, e20))) | (~ (e23 = op2(e23, e23)) & (e22 = op2(e23, e22)) & (e22 = op2(e22, e23))) | (~ (e22 = op2(e22, e22)) & (e22 = op2(e22, e22)) & (e22 = op2(e22, e22))) | (~ (e21 = op2(e21, e21)) & (e22 = op2(e21, e22)) & (e22 = op2(e22, e21))) | (~ (e20 = op2(e20, e20)) & (e22 = op2(e20, e22)) & (e22 = op2(e22, e20))) | (~ (e23 = op2(e23, e23)) & (e21 = op2(e23, e21)) & (e21 = op2(e21, e23))) | (~ (e22 = op2(e22, e22)) & (e21 = op2(e22, e21)) & (e21 = op2(e21, e22))) | (~ (e21 = op2(e21, e21)) & (e21 = op2(e21, e21)) & (e21 = op2(e21, e21))) | (~ (e20 = op2(e20, e20)) & (e21 = op2(e20, e21)) & (e21 = op2(e21, e20))) | (~ (e23 = op2(e23, e23)) & (e20 = op2(e23, e20)) & (e20 = op2(e20, e23))) | (~ (e22 = op2(e22, e22)) & (e20 = op2(e22, e20)) & (e20 = op2(e20, e22))) | (~ (e21 = op2(e21, e21)) & (e20 = op2(e21, e20)) & (e20 = op2(e20, e21))) | (~ (e20 = op2(e20, e20)) & (e20 = op2(e20, e20)) & (e20 = op2(e20, e20)))) & ((e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23))) & ((e23 = op2(e23, e22)) | ~ (e22 = op2(e23, e23))) & ((e23 = op2(e23, e21)) | ~ (e21 = op2(e23, e23))) & ((e23 = op2(e23, e20)) | ~ (e20 = op2(e23, e23))) & ((e22 = op2(e22, e23)) | ~ (e23 = op2(e22, e22))) & ((e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & ((e22 = op2(e22, e21)) | ~ (e21 = op2(e22, e22))) & ((e22 = op2(e22, e20)) | ~ (e20 = op2(e22, e22))) & ((e21 = op2(e21, e23)) | ~ (e23 = op2(e21, e21))) & ((e21 = op2(e21, e22)) | ~ (e22 = op2(e21, e21))) & ((e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & ((e21 = op2(e21, e20)) | ~ (e20 = op2(e21, e21))) & ((e20 = op2(e20, e23)) | ~ (op2(e20, e20) = e23)) & ((e20 = op2(e20, e22)) | ~ (op2(e20, e20) = e22)) & ((e20 = op2(e20, e21)) | ~ (op2(e20, e20) = e21)) & ((e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20))) & ((e23 = op2(e23, e23)) | (e22 = op2(e22, e22)) | (e21 = op2(e21, e21)) | (e20 = op2(e20, e20)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG121+1.p', ax11)).
fof(f1391, plain, (~ spl42_127 | spl42_117), inference(avatar_split_clause, [], [f441, f1070, f1112])).
fof(f1112, plain, (spl42_127 <=> (op2(e20, e20) = e22)), introduced(avatar_definition, [new_symbols(naming, [spl42_127])])).
fof(f441, plain, ((e20 = op2(e20, e22)) | ~ (op2(e20, e20) = e22)), inference(cnf_transformation, [], [f52])).
fof(f1384, plain, (~ spl42_88 | spl42_83), inference(avatar_split_clause, [], [f450, f925, f946])).
fof(f946, plain, (spl42_88 <=> (e23 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_88])])).
fof(f450, plain, ((e22 = op2(e22, e23)) | ~ (e23 = op2(e22, e22))), inference(cnf_transformation, [], [f52])).
fof(f1383, plain, (~ spl42_65 | spl42_80), inference(avatar_split_clause, [], [f451, f912, f849])).
fof(f451, plain, ((e23 = op2(e23, e20)) | ~ (e20 = op2(e23, e23))), inference(cnf_transformation, [], [f52])).
fof(f1381, plain, (~ spl42_67 | spl42_72), inference(avatar_split_clause, [], [f453, f878, f857])).
fof(f857, plain, (spl42_67 <=> (e22 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_67])])).
fof(f453, plain, ((e23 = op2(e23, e22)) | ~ (e22 = op2(e23, e23))), inference(cnf_transformation, [], [f52])).
fof(f1272, plain, (spl42_61 | spl42_42 | spl42_23 | spl42_4), inference(avatar_split_clause, [], [f373, f557, f638, f719, f800])).
fof(f373, plain, ((e13 = op1(e13, e13)) | (e12 = op1(e12, e12)) | (e11 = op1(e11, e11)) | (e10 = op1(e10, e10))), inference(cnf_transformation, [], [f36])).
fof(f36, plain, (((~ (e13 = op1(e13, e13)) & (e13 = op1(e13, e13)) & (e13 = op1(e13, e13))) | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0) & ((e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13))) & ((e13 = op1(e13, e12)) | ~ (e12 = op1(e13, e13))) & ((e13 = op1(e13, e11)) | ~ (e11 = op1(e13, e13))) & ((e13 = op1(e13, e10)) | ~ (e10 = op1(e13, e13))) & ((e12 = op1(e12, e13)) | ~ (e13 = op1(e12, e12))) & ((e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & ((e12 = op1(e12, e11)) | ~ (e11 = op1(e12, e12))) & ((e12 = op1(e12, e10)) | ~ (e10 = op1(e12, e12))) & ((e11 = op1(e11, e13)) | ~ (e13 = op1(e11, e11))) & ((e11 = op1(e11, e12)) | ~ (e12 = op1(e11, e11))) & ((e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & ((e11 = op1(e11, e10)) | ~ (e10 = op1(e11, e11))) & ((e10 = op1(e10, e13)) | ~ (op1(e10, e10) = e13)) & ((e10 = op1(e10, e12)) | ~ (op1(e10, e10) = e12)) & ((e10 = op1(e10, e11)) | ~ (op1(e10, e10) = e11)) & ((e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10))) & ((e13 = op1(e13, e13)) | (e12 = op1(e12, e12)) | (e11 = op1(e11, e11)) | (e10 = op1(e10, e10)))), inference(definition_folding, [], [f10, e35, e34, e33, e32, e31, e30, e29, e28, e27, e26, e25, e24, e23, e22, e21])).
fof(f21, plain, ((~ (e10 = op1(e10, e10)) & (e10 = op1(e10, e10)) & (e10 = op1(e10, e10))) | ~ sP0), inference(usedef, [], [e21])).
fof(e21, plain, (sP0 <=> (~ (e10 = op1(e10, e10)) & (e10 = op1(e10, e10)) & (e10 = op1(e10, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f22, plain, ((~ (e11 = op1(e11, e11)) & (e10 = op1(e11, e10)) & (e10 = op1(e10, e11))) | ~ sP1), inference(usedef, [], [e22])).
fof(e22, plain, (sP1 <=> (~ (e11 = op1(e11, e11)) & (e10 = op1(e11, e10)) & (e10 = op1(e10, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f23, plain, ((~ (e12 = op1(e12, e12)) & (e10 = op1(e12, e10)) & (e10 = op1(e10, e12))) | ~ sP2), inference(usedef, [], [e23])).
fof(e23, plain, (sP2 <=> (~ (e12 = op1(e12, e12)) & (e10 = op1(e12, e10)) & (e10 = op1(e10, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f24, plain, ((~ (e13 = op1(e13, e13)) & (e10 = op1(e13, e10)) & (e10 = op1(e10, e13))) | ~ sP3), inference(usedef, [], [e24])).
fof(e24, plain, (sP3 <=> (~ (e13 = op1(e13, e13)) & (e10 = op1(e13, e10)) & (e10 = op1(e10, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f25, plain, ((~ (e10 = op1(e10, e10)) & (e11 = op1(e10, e11)) & (e11 = op1(e11, e10))) | ~ sP4), inference(usedef, [], [e25])).
fof(e25, plain, (sP4 <=> (~ (e10 = op1(e10, e10)) & (e11 = op1(e10, e11)) & (e11 = op1(e11, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f26, plain, ((~ (e11 = op1(e11, e11)) & (e11 = op1(e11, e11)) & (e11 = op1(e11, e11))) | ~ sP5), inference(usedef, [], [e26])).
fof(e26, plain, (sP5 <=> (~ (e11 = op1(e11, e11)) & (e11 = op1(e11, e11)) & (e11 = op1(e11, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f27, plain, ((~ (e12 = op1(e12, e12)) & (e11 = op1(e12, e11)) & (e11 = op1(e11, e12))) | ~ sP6), inference(usedef, [], [e27])).
fof(e27, plain, (sP6 <=> (~ (e12 = op1(e12, e12)) & (e11 = op1(e12, e11)) & (e11 = op1(e11, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f28, plain, ((~ (e13 = op1(e13, e13)) & (e11 = op1(e13, e11)) & (e11 = op1(e11, e13))) | ~ sP7), inference(usedef, [], [e28])).
fof(e28, plain, (sP7 <=> (~ (e13 = op1(e13, e13)) & (e11 = op1(e13, e11)) & (e11 = op1(e11, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f29, plain, ((~ (e10 = op1(e10, e10)) & (e12 = op1(e10, e12)) & (e12 = op1(e12, e10))) | ~ sP8), inference(usedef, [], [e29])).
fof(e29, plain, (sP8 <=> (~ (e10 = op1(e10, e10)) & (e12 = op1(e10, e12)) & (e12 = op1(e12, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f30, plain, ((~ (e11 = op1(e11, e11)) & (e12 = op1(e11, e12)) & (e12 = op1(e12, e11))) | ~ sP9), inference(usedef, [], [e30])).
fof(e30, plain, (sP9 <=> (~ (e11 = op1(e11, e11)) & (e12 = op1(e11, e12)) & (e12 = op1(e12, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f31, plain, ((~ (e12 = op1(e12, e12)) & (e12 = op1(e12, e12)) & (e12 = op1(e12, e12))) | ~ sP10), inference(usedef, [], [e31])).
fof(e31, plain, (sP10 <=> (~ (e12 = op1(e12, e12)) & (e12 = op1(e12, e12)) & (e12 = op1(e12, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f32, plain, ((~ (e13 = op1(e13, e13)) & (e12 = op1(e13, e12)) & (e12 = op1(e12, e13))) | ~ sP11), inference(usedef, [], [e32])).
fof(e32, plain, (sP11 <=> (~ (e13 = op1(e13, e13)) & (e12 = op1(e13, e12)) & (e12 = op1(e12, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f33, plain, ((~ (e10 = op1(e10, e10)) & (e13 = op1(e10, e13)) & (e13 = op1(e13, e10))) | ~ sP12), inference(usedef, [], [e33])).
fof(e33, plain, (sP12 <=> (~ (e10 = op1(e10, e10)) & (e13 = op1(e10, e13)) & (e13 = op1(e13, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f34, plain, ((~ (e11 = op1(e11, e11)) & (e13 = op1(e11, e13)) & (e13 = op1(e13, e11))) | ~ sP13), inference(usedef, [], [e34])).
fof(e34, plain, (sP13 <=> (~ (e11 = op1(e11, e11)) & (e13 = op1(e11, e13)) & (e13 = op1(e13, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f35, plain, ((~ (e12 = op1(e12, e12)) & (e13 = op1(e12, e13)) & (e13 = op1(e13, e12))) | ~ sP14), inference(usedef, [], [e35])).
fof(e35, plain, (sP14 <=> (~ (e12 = op1(e12, e12)) & (e13 = op1(e12, e13)) & (e13 = op1(e13, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f10, plain, (((~ (e13 = op1(e13, e13)) & (e13 = op1(e13, e13)) & (e13 = op1(e13, e13))) | (~ (e12 = op1(e12, e12)) & (e13 = op1(e12, e13)) & (e13 = op1(e13, e12))) | (~ (e11 = op1(e11, e11)) & (e13 = op1(e11, e13)) & (e13 = op1(e13, e11))) | (~ (e10 = op1(e10, e10)) & (e13 = op1(e10, e13)) & (e13 = op1(e13, e10))) | (~ (e13 = op1(e13, e13)) & (e12 = op1(e13, e12)) & (e12 = op1(e12, e13))) | (~ (e12 = op1(e12, e12)) & (e12 = op1(e12, e12)) & (e12 = op1(e12, e12))) | (~ (e11 = op1(e11, e11)) & (e12 = op1(e11, e12)) & (e12 = op1(e12, e11))) | (~ (e10 = op1(e10, e10)) & (e12 = op1(e10, e12)) & (e12 = op1(e12, e10))) | (~ (e13 = op1(e13, e13)) & (e11 = op1(e13, e11)) & (e11 = op1(e11, e13))) | (~ (e12 = op1(e12, e12)) & (e11 = op1(e12, e11)) & (e11 = op1(e11, e12))) | (~ (e11 = op1(e11, e11)) & (e11 = op1(e11, e11)) & (e11 = op1(e11, e11))) | (~ (e10 = op1(e10, e10)) & (e11 = op1(e10, e11)) & (e11 = op1(e11, e10))) | (~ (e13 = op1(e13, e13)) & (e10 = op1(e13, e10)) & (e10 = op1(e10, e13))) | (~ (e12 = op1(e12, e12)) & (e10 = op1(e12, e10)) & (e10 = op1(e10, e12))) | (~ (e11 = op1(e11, e11)) & (e10 = op1(e11, e10)) & (e10 = op1(e10, e11))) | (~ (e10 = op1(e10, e10)) & (e10 = op1(e10, e10)) & (e10 = op1(e10, e10)))) & ((e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13))) & ((e13 = op1(e13, e12)) | ~ (e12 = op1(e13, e13))) & ((e13 = op1(e13, e11)) | ~ (e11 = op1(e13, e13))) & ((e13 = op1(e13, e10)) | ~ (e10 = op1(e13, e13))) & ((e12 = op1(e12, e13)) | ~ (e13 = op1(e12, e12))) & ((e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & ((e12 = op1(e12, e11)) | ~ (e11 = op1(e12, e12))) & ((e12 = op1(e12, e10)) | ~ (e10 = op1(e12, e12))) & ((e11 = op1(e11, e13)) | ~ (e13 = op1(e11, e11))) & ((e11 = op1(e11, e12)) | ~ (e12 = op1(e11, e11))) & ((e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & ((e11 = op1(e11, e10)) | ~ (e10 = op1(e11, e11))) & ((e10 = op1(e10, e13)) | ~ (op1(e10, e10) = e13)) & ((e10 = op1(e10, e12)) | ~ (op1(e10, e10) = e12)) & ((e10 = op1(e10, e11)) | ~ (op1(e10, e10) = e11)) & ((e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10))) & ((e13 = op1(e13, e13)) | (e12 = op1(e12, e12)) | (e11 = op1(e11, e11)) | (e10 = op1(e10, e10)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG121+1.p', ax10)).
fof(f1271, plain, (~ spl42_62 | spl42_57), inference(avatar_split_clause, [], [f375, f783, f804])).
fof(f375, plain, ((e10 = op1(e10, e11)) | ~ (op1(e10, e10) = e11)), inference(cnf_transformation, [], [f36])).
fof(f1270, plain, (~ spl42_63 | spl42_53), inference(avatar_split_clause, [], [f376, f766, f808])).
fof(f808, plain, (spl42_63 <=> (op1(e10, e10) = e12)), introduced(avatar_definition, [new_symbols(naming, [spl42_63])])).
fof(f376, plain, ((e10 = op1(e10, e12)) | ~ (op1(e10, e10) = e12)), inference(cnf_transformation, [], [f36])).
fof(f1269, plain, (~ spl42_64 | spl42_49), inference(avatar_split_clause, [], [f377, f749, f812])).
fof(f812, plain, (spl42_64 <=> (op1(e10, e10) = e13)), introduced(avatar_definition, [new_symbols(naming, [spl42_64])])).
fof(f377, plain, ((e10 = op1(e10, e13)) | ~ (op1(e10, e10) = e13)), inference(cnf_transformation, [], [f36])).
fof(f1266, plain, (~ spl42_44 | spl42_34), inference(avatar_split_clause, [], [f381, f685, f727])).
fof(f727, plain, (spl42_44 <=> (e13 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_44])])).
fof(f381, plain, ((e11 = op1(e11, e13)) | ~ (e13 = op1(e11, e11))), inference(cnf_transformation, [], [f36])).
fof(f1264, plain, (~ spl42_22 | spl42_27), inference(avatar_split_clause, [], [f383, f655, f634])).
fof(f634, plain, (spl42_22 <=> (e11 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_22])])).
fof(f383, plain, ((e12 = op1(e12, e11)) | ~ (e11 = op1(e12, e12))), inference(cnf_transformation, [], [f36])).
fof(f1262, plain, (~ spl42_1 | spl42_16), inference(avatar_split_clause, [], [f386, f608, f545])).
fof(f386, plain, ((e13 = op1(e13, e10)) | ~ (e10 = op1(e13, e13))), inference(cnf_transformation, [], [f36])).
fof(f1146, plain, (spl42_127 | spl42_111 | spl42_95 | spl42_79), inference(avatar_split_clause, [], [f177, f908, f976, f1044, f1112])).
fof(f177, plain, ((e22 = op2(e23, e20)) | (e22 = op2(e22, e20)) | (e22 = op2(e21, e20)) | (op2(e20, e20) = e22)), inference(cnf_transformation, [], [f4])).
fof(f4, plain, (((e23 = op2(e23, e23)) | (e23 = op2(e22, e23)) | (e23 = op2(e21, e23)) | (e23 = op2(e20, e23))) & ((e23 = op2(e23, e23)) | (e23 = op2(e23, e22)) | (e23 = op2(e23, e21)) | (e23 = op2(e23, e20))) & ((e22 = op2(e23, e23)) | (e22 = op2(e22, e23)) | (e22 = op2(e21, e23)) | (e22 = op2(e20, e23))) & ((e22 = op2(e23, e23)) | (e22 = op2(e23, e22)) | (e22 = op2(e23, e21)) | (e22 = op2(e23, e20))) & ((e21 = op2(e23, e23)) | (e21 = op2(e22, e23)) | (e21 = op2(e21, e23)) | (e21 = op2(e20, e23))) & ((e21 = op2(e23, e23)) | (e21 = op2(e23, e22)) | (e21 = op2(e23, e21)) | (e21 = op2(e23, e20))) & ((e20 = op2(e23, e23)) | (e20 = op2(e22, e23)) | (e20 = op2(e21, e23)) | (e20 = op2(e20, e23))) & ((e20 = op2(e23, e23)) | (e20 = op2(e23, e22)) | (e20 = op2(e23, e21)) | (e20 = op2(e23, e20))) & ((e23 = op2(e23, e22)) | (e23 = op2(e22, e22)) | (e23 = op2(e21, e22)) | (e23 = op2(e20, e22))) & ((e23 = op2(e22, e23)) | (e23 = op2(e22, e22)) | (e23 = op2(e22, e21)) | (e23 = op2(e22, e20))) & ((e22 = op2(e23, e22)) | (e22 = op2(e22, e22)) | (e22 = op2(e21, e22)) | (e22 = op2(e20, e22))) & ((e22 = op2(e22, e23)) | (e22 = op2(e22, e22)) | (e22 = op2(e22, e21)) | (e22 = op2(e22, e20))) & ((e21 = op2(e23, e22)) | (e21 = op2(e22, e22)) | (e21 = op2(e21, e22)) | (e21 = op2(e20, e22))) & ((e21 = op2(e22, e23)) | (e21 = op2(e22, e22)) | (e21 = op2(e22, e21)) | (e21 = op2(e22, e20))) & ((e20 = op2(e23, e22)) | (e20 = op2(e22, e22)) | (e20 = op2(e21, e22)) | (e20 = op2(e20, e22))) & ((e20 = op2(e22, e23)) | (e20 = op2(e22, e22)) | (e20 = op2(e22, e21)) | (e20 = op2(e22, e20))) & ((e23 = op2(e23, e21)) | (e23 = op2(e22, e21)) | (e23 = op2(e21, e21)) | (e23 = op2(e20, e21))) & ((e23 = op2(e21, e23)) | (e23 = op2(e21, e22)) | (e23 = op2(e21, e21)) | (e23 = op2(e21, e20))) & ((e22 = op2(e23, e21)) | (e22 = op2(e22, e21)) | (e22 = op2(e21, e21)) | (e22 = op2(e20, e21))) & ((e22 = op2(e21, e23)) | (e22 = op2(e21, e22)) | (e22 = op2(e21, e21)) | (e22 = op2(e21, e20))) & ((e21 = op2(e23, e21)) | (e21 = op2(e22, e21)) | (e21 = op2(e21, e21)) | (e21 = op2(e20, e21))) & ((e21 = op2(e21, e23)) | (e21 = op2(e21, e22)) | (e21 = op2(e21, e21)) | (e21 = op2(e21, e20))) & ((e20 = op2(e23, e21)) | (e20 = op2(e22, e21)) | (e20 = op2(e21, e21)) | (e20 = op2(e20, e21))) & ((e20 = op2(e21, e23)) | (e20 = op2(e21, e22)) | (e20 = op2(e21, e21)) | (e20 = op2(e21, e20))) & ((e23 = op2(e23, e20)) | (e23 = op2(e22, e20)) | (e23 = op2(e21, e20)) | (op2(e20, e20) = e23)) & ((e23 = op2(e20, e23)) | (e23 = op2(e20, e22)) | (e23 = op2(e20, e21)) | (op2(e20, e20) = e23)) & ((e22 = op2(e23, e20)) | (e22 = op2(e22, e20)) | (e22 = op2(e21, e20)) | (op2(e20, e20) = e22)) & ((e22 = op2(e20, e23)) | (e22 = op2(e20, e22)) | (e22 = op2(e20, e21)) | (op2(e20, e20) = e22)) & ((e21 = op2(e23, e20)) | (e21 = op2(e22, e20)) | (e21 = op2(e21, e20)) | (op2(e20, e20) = e21)) & ((e21 = op2(e20, e23)) | (e21 = op2(e20, e22)) | (e21 = op2(e20, e21)) | (op2(e20, e20) = e21)) & ((e20 = op2(e23, e20)) | (e20 = op2(e22, e20)) | (e20 = op2(e21, e20)) | (e20 = op2(e20, e20))) & ((e20 = op2(e20, e23)) | (e20 = op2(e20, e22)) | (e20 = op2(e20, e21)) | (e20 = op2(e20, e20)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG121+1.p', ax4)).
fof(f1140, plain, (spl42_122 | spl42_106 | spl42_90 | spl42_74), inference(avatar_split_clause, [], [f183, f887, f955, f1023, f1091])).
fof(f183, plain, ((e21 = op2(e23, e21)) | (e21 = op2(e22, e21)) | (e21 = op2(e21, e21)) | (e21 = op2(e20, e21))), inference(cnf_transformation, [], [f4])).
fof(f1135, plain, (spl42_93 | spl42_89 | spl42_85 | spl42_81), inference(avatar_split_clause, [], [f188, f917, f934, f951, f968])).
fof(f188, plain, ((e20 = op2(e22, e23)) | (e20 = op2(e22, e22)) | (e20 = op2(e22, e21)) | (e20 = op2(e22, e20))), inference(cnf_transformation, [], [f4])).
fof(f1129, plain, (spl42_96 | spl42_92 | spl42_88 | spl42_84), inference(avatar_split_clause, [], [f194, f929, f946, f963, f980])).
fof(f194, plain, ((e23 = op2(e22, e23)) | (e23 = op2(e22, e22)) | (e23 = op2(e22, e21)) | (e23 = op2(e22, e20))), inference(cnf_transformation, [], [f4])).
fof(f1128, plain, (spl42_120 | spl42_104 | spl42_88 | spl42_72), inference(avatar_split_clause, [], [f195, f878, f946, f1014, f1082])).
fof(f195, plain, ((e23 = op2(e23, e22)) | (e23 = op2(e22, e22)) | (e23 = op2(e21, e22)) | (e23 = op2(e20, e22))), inference(cnf_transformation, [], [f4])).
fof(f1123, plain, (spl42_79 | spl42_75 | spl42_71 | spl42_67), inference(avatar_split_clause, [], [f200, f857, f874, f891, f908])).
fof(f200, plain, ((e22 = op2(e23, e23)) | (e22 = op2(e23, e22)) | (e22 = op2(e23, e21)) | (e22 = op2(e23, e20))), inference(cnf_transformation, [], [f4])).
fof(f1085, plain, (spl42_117 | spl42_118 | spl42_119 | spl42_120), inference(avatar_split_clause, [], [f158, f1082, f1078, f1074, f1070])).
fof(f158, plain, ((e23 = op2(e20, e22)) | (e22 = op2(e20, e22)) | (e21 = op2(e20, e22)) | (e20 = op2(e20, e22))), inference(cnf_transformation, [], [f3])).
fof(f3, plain, (((e23 = op2(e23, e23)) | (e22 = op2(e23, e23)) | (e21 = op2(e23, e23)) | (e20 = op2(e23, e23))) & ((e23 = op2(e23, e22)) | (e22 = op2(e23, e22)) | (e21 = op2(e23, e22)) | (e20 = op2(e23, e22))) & ((e23 = op2(e23, e21)) | (e22 = op2(e23, e21)) | (e21 = op2(e23, e21)) | (e20 = op2(e23, e21))) & ((e23 = op2(e23, e20)) | (e22 = op2(e23, e20)) | (e21 = op2(e23, e20)) | (e20 = op2(e23, e20))) & ((e23 = op2(e22, e23)) | (e22 = op2(e22, e23)) | (e21 = op2(e22, e23)) | (e20 = op2(e22, e23))) & ((e23 = op2(e22, e22)) | (e22 = op2(e22, e22)) | (e21 = op2(e22, e22)) | (e20 = op2(e22, e22))) & ((e23 = op2(e22, e21)) | (e22 = op2(e22, e21)) | (e21 = op2(e22, e21)) | (e20 = op2(e22, e21))) & ((e23 = op2(e22, e20)) | (e22 = op2(e22, e20)) | (e21 = op2(e22, e20)) | (e20 = op2(e22, e20))) & ((e23 = op2(e21, e23)) | (e22 = op2(e21, e23)) | (e21 = op2(e21, e23)) | (e20 = op2(e21, e23))) & ((e23 = op2(e21, e22)) | (e22 = op2(e21, e22)) | (e21 = op2(e21, e22)) | (e20 = op2(e21, e22))) & ((e23 = op2(e21, e21)) | (e22 = op2(e21, e21)) | (e21 = op2(e21, e21)) | (e20 = op2(e21, e21))) & ((e23 = op2(e21, e20)) | (e22 = op2(e21, e20)) | (e21 = op2(e21, e20)) | (e20 = op2(e21, e20))) & ((e23 = op2(e20, e23)) | (e22 = op2(e20, e23)) | (e21 = op2(e20, e23)) | (e20 = op2(e20, e23))) & ((e23 = op2(e20, e22)) | (e22 = op2(e20, e22)) | (e21 = op2(e20, e22)) | (e20 = op2(e20, e22))) & ((e23 = op2(e20, e21)) | (e22 = op2(e20, e21)) | (e21 = op2(e20, e21)) | (e20 = op2(e20, e21))) & ((op2(e20, e20) = e23) | (op2(e20, e20) = e22) | (op2(e20, e20) = e21) | (e20 = op2(e20, e20)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG121+1.p', ax3)).
fof(f1068, plain, (spl42_113 | spl42_114 | spl42_115 | spl42_116), inference(avatar_split_clause, [], [f159, f1065, f1061, f1057, f1053])).
fof(f159, plain, ((e23 = op2(e20, e23)) | (e22 = op2(e20, e23)) | (e21 = op2(e20, e23)) | (e20 = op2(e20, e23))), inference(cnf_transformation, [], [f3])).
fof(f1051, plain, (spl42_109 | spl42_110 | spl42_111 | spl42_112), inference(avatar_split_clause, [], [f160, f1048, f1044, f1040, f1036])).
fof(f160, plain, ((e23 = op2(e21, e20)) | (e22 = op2(e21, e20)) | (e21 = op2(e21, e20)) | (e20 = op2(e21, e20))), inference(cnf_transformation, [], [f3])).
fof(f932, plain, (spl42_81 | spl42_82 | spl42_83 | spl42_84), inference(avatar_split_clause, [], [f167, f929, f925, f921, f917])).
fof(f167, plain, ((e23 = op2(e22, e23)) | (e22 = op2(e22, e23)) | (e21 = op2(e22, e23)) | (e20 = op2(e22, e23))), inference(cnf_transformation, [], [f3])).
fof(f881, plain, (spl42_69 | spl42_70 | spl42_71 | spl42_72), inference(avatar_split_clause, [], [f170, f878, f874, f870, f866])).
fof(f170, plain, ((e23 = op2(e23, e22)) | (e22 = op2(e23, e22)) | (e21 = op2(e23, e22)) | (e20 = op2(e23, e22))), inference(cnf_transformation, [], [f3])).
fof(f842, plain, (spl42_63 | spl42_47 | spl42_31 | spl42_15), inference(avatar_split_clause, [], [f129, f604, f672, f740, f808])).
fof(f129, plain, ((e12 = op1(e13, e10)) | (e12 = op1(e12, e10)) | (e12 = op1(e11, e10)) | (op1(e10, e10) = e12)), inference(cnf_transformation, [], [f2])).
fof(f2, plain, (((e13 = op1(e13, e13)) | (e13 = op1(e12, e13)) | (e13 = op1(e11, e13)) | (e13 = op1(e10, e13))) & ((e13 = op1(e13, e13)) | (e13 = op1(e13, e12)) | (e13 = op1(e13, e11)) | (e13 = op1(e13, e10))) & ((e12 = op1(e13, e13)) | (e12 = op1(e12, e13)) | (e12 = op1(e11, e13)) | (e12 = op1(e10, e13))) & ((e12 = op1(e13, e13)) | (e12 = op1(e13, e12)) | (e12 = op1(e13, e11)) | (e12 = op1(e13, e10))) & ((e11 = op1(e13, e13)) | (e11 = op1(e12, e13)) | (e11 = op1(e11, e13)) | (e11 = op1(e10, e13))) & ((e11 = op1(e13, e13)) | (e11 = op1(e13, e12)) | (e11 = op1(e13, e11)) | (e11 = op1(e13, e10))) & ((e10 = op1(e13, e13)) | (e10 = op1(e12, e13)) | (e10 = op1(e11, e13)) | (e10 = op1(e10, e13))) & ((e10 = op1(e13, e13)) | (e10 = op1(e13, e12)) | (e10 = op1(e13, e11)) | (e10 = op1(e13, e10))) & ((e13 = op1(e13, e12)) | (e13 = op1(e12, e12)) | (e13 = op1(e11, e12)) | (e13 = op1(e10, e12))) & ((e13 = op1(e12, e13)) | (e13 = op1(e12, e12)) | (e13 = op1(e12, e11)) | (e13 = op1(e12, e10))) & ((e12 = op1(e13, e12)) | (e12 = op1(e12, e12)) | (e12 = op1(e11, e12)) | (e12 = op1(e10, e12))) & ((e12 = op1(e12, e13)) | (e12 = op1(e12, e12)) | (e12 = op1(e12, e11)) | (e12 = op1(e12, e10))) & ((e11 = op1(e13, e12)) | (e11 = op1(e12, e12)) | (e11 = op1(e11, e12)) | (e11 = op1(e10, e12))) & ((e11 = op1(e12, e13)) | (e11 = op1(e12, e12)) | (e11 = op1(e12, e11)) | (e11 = op1(e12, e10))) & ((e10 = op1(e13, e12)) | (e10 = op1(e12, e12)) | (e10 = op1(e11, e12)) | (e10 = op1(e10, e12))) & ((e10 = op1(e12, e13)) | (e10 = op1(e12, e12)) | (e10 = op1(e12, e11)) | (e10 = op1(e12, e10))) & ((e13 = op1(e13, e11)) | (e13 = op1(e12, e11)) | (e13 = op1(e11, e11)) | (e13 = op1(e10, e11))) & ((e13 = op1(e11, e13)) | (e13 = op1(e11, e12)) | (e13 = op1(e11, e11)) | (e13 = op1(e11, e10))) & ((e12 = op1(e13, e11)) | (e12 = op1(e12, e11)) | (e12 = op1(e11, e11)) | (e12 = op1(e10, e11))) & ((e12 = op1(e11, e13)) | (e12 = op1(e11, e12)) | (e12 = op1(e11, e11)) | (e12 = op1(e11, e10))) & ((e11 = op1(e13, e11)) | (e11 = op1(e12, e11)) | (e11 = op1(e11, e11)) | (e11 = op1(e10, e11))) & ((e11 = op1(e11, e13)) | (e11 = op1(e11, e12)) | (e11 = op1(e11, e11)) | (e11 = op1(e11, e10))) & ((e10 = op1(e13, e11)) | (e10 = op1(e12, e11)) | (e10 = op1(e11, e11)) | (e10 = op1(e10, e11))) & ((e10 = op1(e11, e13)) | (e10 = op1(e11, e12)) | (e10 = op1(e11, e11)) | (e10 = op1(e11, e10))) & ((e13 = op1(e13, e10)) | (e13 = op1(e12, e10)) | (e13 = op1(e11, e10)) | (op1(e10, e10) = e13)) & ((e13 = op1(e10, e13)) | (e13 = op1(e10, e12)) | (e13 = op1(e10, e11)) | (op1(e10, e10) = e13)) & ((e12 = op1(e13, e10)) | (e12 = op1(e12, e10)) | (e12 = op1(e11, e10)) | (op1(e10, e10) = e12)) & ((e12 = op1(e10, e13)) | (e12 = op1(e10, e12)) | (e12 = op1(e10, e11)) | (op1(e10, e10) = e12)) & ((e11 = op1(e13, e10)) | (e11 = op1(e12, e10)) | (e11 = op1(e11, e10)) | (op1(e10, e10) = e11)) & ((e11 = op1(e10, e13)) | (e11 = op1(e10, e12)) | (e11 = op1(e10, e11)) | (op1(e10, e10) = e11)) & ((e10 = op1(e13, e10)) | (e10 = op1(e12, e10)) | (e10 = op1(e11, e10)) | (e10 = op1(e10, e10))) & ((e10 = op1(e10, e13)) | (e10 = op1(e10, e12)) | (e10 = op1(e10, e11)) | (e10 = op1(e10, e10)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG121+1.p', ax2)).
fof(f841, plain, (spl42_64 | spl42_60 | spl42_56 | spl42_52), inference(avatar_split_clause, [], [f130, f761, f778, f795, f812])).
fof(f130, plain, ((e13 = op1(e10, e13)) | (e13 = op1(e10, e12)) | (e13 = op1(e10, e11)) | (op1(e10, e10) = e13)), inference(cnf_transformation, [], [f2])).
fof(f833, plain, (spl42_48 | spl42_44 | spl42_40 | spl42_36), inference(avatar_split_clause, [], [f138, f693, f710, f727, f744])).
fof(f138, plain, ((e13 = op1(e11, e13)) | (e13 = op1(e11, e12)) | (e13 = op1(e11, e11)) | (e13 = op1(e11, e10))), inference(cnf_transformation, [], [f2])).
fof(f832, plain, (spl42_60 | spl42_44 | spl42_28 | spl42_12), inference(avatar_split_clause, [], [f139, f591, f659, f727, f795])).
fof(f139, plain, ((e13 = op1(e13, e11)) | (e13 = op1(e12, e11)) | (e13 = op1(e11, e11)) | (e13 = op1(e10, e11))), inference(cnf_transformation, [], [f2])).
fof(f831, plain, (spl42_29 | spl42_25 | spl42_21 | spl42_17), inference(avatar_split_clause, [], [f140, f613, f630, f647, f664])).
fof(f140, plain, ((e10 = op1(e12, e13)) | (e10 = op1(e12, e12)) | (e10 = op1(e12, e11)) | (e10 = op1(e12, e10))), inference(cnf_transformation, [], [f2])).
fof(f829, plain, (spl42_30 | spl42_26 | spl42_22 | spl42_18), inference(avatar_split_clause, [], [f142, f617, f634, f651, f668])).
fof(f142, plain, ((e11 = op1(e12, e13)) | (e11 = op1(e12, e12)) | (e11 = op1(e12, e11)) | (e11 = op1(e12, e10))), inference(cnf_transformation, [], [f2])).
fof(f828, plain, (spl42_54 | spl42_38 | spl42_22 | spl42_6), inference(avatar_split_clause, [], [f143, f566, f634, f702, f770])).
fof(f143, plain, ((e11 = op1(e13, e12)) | (e11 = op1(e12, e12)) | (e11 = op1(e11, e12)) | (e11 = op1(e10, e12))), inference(cnf_transformation, [], [f2])).
fof(f826, plain, (spl42_55 | spl42_39 | spl42_23 | spl42_7), inference(avatar_split_clause, [], [f145, f570, f638, f706, f774])).
fof(f145, plain, ((e12 = op1(e13, e12)) | (e12 = op1(e12, e12)) | (e12 = op1(e11, e12)) | (e12 = op1(e10, e12))), inference(cnf_transformation, [], [f2])).
fof(f819, plain, (spl42_15 | spl42_11 | spl42_7 | spl42_3), inference(avatar_split_clause, [], [f152, f553, f570, f587, f604])).
fof(f152, plain, ((e12 = op1(e13, e13)) | (e12 = op1(e13, e12)) | (e12 = op1(e13, e11)) | (e12 = op1(e13, e10))), inference(cnf_transformation, [], [f2])).
fof(f747, plain, (spl42_45 | spl42_46 | spl42_47 | spl42_48), inference(avatar_split_clause, [], [f112, f744, f740, f736, f732])).
fof(f112, plain, ((e13 = op1(e11, e10)) | (e12 = op1(e11, e10)) | (e11 = op1(e11, e10)) | (e10 = op1(e11, e10))), inference(cnf_transformation, [], [f1])).
fof(f1, plain, (((e13 = op1(e13, e13)) | (e12 = op1(e13, e13)) | (e11 = op1(e13, e13)) | (e10 = op1(e13, e13))) & ((e13 = op1(e13, e12)) | (e12 = op1(e13, e12)) | (e11 = op1(e13, e12)) | (e10 = op1(e13, e12))) & ((e13 = op1(e13, e11)) | (e12 = op1(e13, e11)) | (e11 = op1(e13, e11)) | (e10 = op1(e13, e11))) & ((e13 = op1(e13, e10)) | (e12 = op1(e13, e10)) | (e11 = op1(e13, e10)) | (e10 = op1(e13, e10))) & ((e13 = op1(e12, e13)) | (e12 = op1(e12, e13)) | (e11 = op1(e12, e13)) | (e10 = op1(e12, e13))) & ((e13 = op1(e12, e12)) | (e12 = op1(e12, e12)) | (e11 = op1(e12, e12)) | (e10 = op1(e12, e12))) & ((e13 = op1(e12, e11)) | (e12 = op1(e12, e11)) | (e11 = op1(e12, e11)) | (e10 = op1(e12, e11))) & ((e13 = op1(e12, e10)) | (e12 = op1(e12, e10)) | (e11 = op1(e12, e10)) | (e10 = op1(e12, e10))) & ((e13 = op1(e11, e13)) | (e12 = op1(e11, e13)) | (e11 = op1(e11, e13)) | (e10 = op1(e11, e13))) & ((e13 = op1(e11, e12)) | (e12 = op1(e11, e12)) | (e11 = op1(e11, e12)) | (e10 = op1(e11, e12))) & ((e13 = op1(e11, e11)) | (e12 = op1(e11, e11)) | (e11 = op1(e11, e11)) | (e10 = op1(e11, e11))) & ((e13 = op1(e11, e10)) | (e12 = op1(e11, e10)) | (e11 = op1(e11, e10)) | (e10 = op1(e11, e10))) & ((e13 = op1(e10, e13)) | (e12 = op1(e10, e13)) | (e11 = op1(e10, e13)) | (e10 = op1(e10, e13))) & ((e13 = op1(e10, e12)) | (e12 = op1(e10, e12)) | (e11 = op1(e10, e12)) | (e10 = op1(e10, e12))) & ((e13 = op1(e10, e11)) | (e12 = op1(e10, e11)) | (e11 = op1(e10, e11)) | (e10 = op1(e10, e11))) & ((op1(e10, e10) = e13) | (op1(e10, e10) = e12) | (op1(e10, e10) = e11) | (e10 = op1(e10, e10)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG121+1.p', ax1)).