fof(f5253, plain, $false, inference(avatar_sat_refutation, [], [f1203, f1373, f1442, f1449, f1454, f1457, f1462, f1466, f1468, f1469, f1677, f1753, f1754, f1755, f1758, f1759, f1761, f1764, f1771, f1772, f2378, f2379, f2381, f2982, f2983, f2985, f2986, f2988, f3041, f3055, f3271, f3661, f3662, f3693, f3711, f3716, f3717, f3718, f3746, f3749, f3808, f3811, f3821, f3822, f3864, f3872, f3878, f3889, f3904, f3935, f3967, f3975, f3977, f3981, f3988, f4023, f4025, f4064, f4072, f4087, f4090, f4092, f4126, f4139, f4143, f4156, f4201, f4214, f4227, f4267, f4296, f4306, f4319, f4329, f4358, f4380, f4383, f4394, f4411, f4417, f4461, f4469, f4485, f4507, f4510, f4571, f4593, f4596, f4643, f4685, f4686, f4687, f4700, f4708, f4718, f4774, f4808, f4825, f4829, f4858, f4870, f4902, f4932, f4965, f5000, f5030, f5058, f5087, f5118, f5149, f5174, f5199, f5229, f5252])).
fof(f5252, plain, (~ spl144_54 | ~ spl144_118 | ~ spl144_294 | spl144_334), inference(avatar_contradiction_clause, [], [f5251])).
fof(f5251, plain, ($false | (~ spl144_54 | ~ spl144_118 | ~ spl144_294 | spl144_334)), inference(subsumption_resolution, [], [f5250, f1702])).
fof(f1702, plain, ((e21 = op2(e20, e22)) | ~ spl144_118), inference(avatar_component_clause, [], [f1700])).
fof(f1700, plain, (spl144_118 <=> (e21 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl144_118])])).
fof(f5250, plain, (~ (e21 = op2(e20, e22)) | (~ spl144_54 | ~ spl144_294 | spl144_334)), inference(forward_demodulation, [], [f5249, f3039])).
fof(f3039, plain, ((e21 = h4(e11)) | ~ spl144_294), inference(avatar_component_clause, [], [f3038])).
fof(f3038, plain, (spl144_294 <=> (e21 = h4(e11))), introduced(avatar_definition, [new_symbols(naming, [spl144_294])])).
fof(f5249, plain, (~ (op2(e20, e22) = h4(e11)) | (~ spl144_54 | spl144_334)), inference(forward_demodulation, [], [f3262, f1398])).
fof(f1398, plain, ((e11 = op1(e10, e12)) | ~ spl144_54), inference(avatar_component_clause, [], [f1396])).
fof(f1396, plain, (spl144_54 <=> (e11 = op1(e10, e12))), introduced(avatar_definition, [new_symbols(naming, [spl144_54])])).
fof(f3262, plain, (~ (op2(e20, e22) = h4(op1(e10, e12))) | spl144_334), inference(avatar_component_clause, [], [f3260])).
fof(f3260, plain, (spl144_334 <=> (op2(e20, e22) = h4(op1(e10, e12)))), introduced(avatar_definition, [new_symbols(naming, [spl144_334])])).
fof(f5229, plain, (~ spl144_61 | ~ spl144_125 | spl144_336), inference(avatar_contradiction_clause, [], [f5228])).
fof(f5228, plain, ($false | (~ spl144_61 | ~ spl144_125 | spl144_336)), inference(subsumption_resolution, [], [f5227, f4435])).
fof(f4435, plain, ((e20 = h1(e11)) | ~ spl144_125), inference(backward_demodulation, [], [f1092, f1732])).
fof(f1732, plain, ((e20 = op2(e20, e20)) | ~ spl144_125), inference(avatar_component_clause, [], [f1730])).
fof(f1730, plain, (spl144_125 <=> (e20 = op2(e20, e20))), introduced(avatar_definition, [new_symbols(naming, [spl144_125])])).
fof(f1092, plain, (op2(e20, e20) = h1(e11)), inference(cnf_transformation, [], [f14])).
fof(f14, plain, ((op2(op2(e20, e20), e20) = h1(e12)) & (op2(e20, e20) = h1(e11)) & (h1(e10) = op2(op2(e20, e20), op2(op2(e20, e20), e20))) & (e20 = h1(e13))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG115+1.p', ax14)).
fof(f5227, plain, (~ (e20 = h1(e11)) | (~ spl144_61 | spl144_336)), inference(forward_demodulation, [], [f5226, f3028])).
fof(f3028, plain, (e20 = h4(e10)), inference(forward_demodulation, [], [f3027, f3025])).
fof(f3025, plain, (e20 = op2(h4(e11), e22)), inference(backward_demodulation, [], [f2989, f1104])).
fof(f1104, plain, (op2(e23, e23) = h4(e11)), inference(cnf_transformation, [], [f17])).
fof(f17, plain, ((op2(op2(e23, e23), e23) = h4(e12)) & (op2(e23, e23) = h4(e11)) & (op2(op2(e23, e23), op2(op2(e23, e23), e23)) = h4(e10)) & (e23 = h4(e13))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG115+1.p', ax17)).
fof(f2989, plain, (e20 = op2(op2(e23, e23), e22)), inference(forward_demodulation, [], [f1087, f1089])).
fof(f1089, plain, (e22 = op2(op2(e23, e23), e23)), inference(cnf_transformation, [], [f13])).
fof(f13, plain, ((e22 = op2(op2(e23, e23), e23)) & (e21 = op2(e23, e23)) & (e20 = op2(op2(e23, e23), op2(op2(e23, e23), e23)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG115+1.p', ax13)).
fof(f1087, plain, (e20 = op2(op2(e23, e23), op2(op2(e23, e23), e23))), inference(cnf_transformation, [], [f13])).
fof(f3027, plain, (h4(e10) = op2(h4(e11), e22)), inference(forward_demodulation, [], [f3026, f3024])).
fof(f3024, plain, (e22 = op2(h4(e11), e23)), inference(backward_demodulation, [], [f1089, f1104])).
fof(f3026, plain, (h4(e10) = op2(h4(e11), op2(h4(e11), e23))), inference(forward_demodulation, [], [f1103, f1104])).
fof(f1103, plain, (op2(op2(e23, e23), op2(op2(e23, e23), e23)) = h4(e10)), inference(cnf_transformation, [], [f17])).
fof(f5226, plain, (~ (h1(e11) = h4(e10)) | (~ spl144_61 | spl144_336)), inference(forward_demodulation, [], [f3270, f1428])).
fof(f1428, plain, ((e10 = op1(e10, e10)) | ~ spl144_61), inference(avatar_component_clause, [], [f1426])).
fof(f1426, plain, (spl144_61 <=> (e10 = op1(e10, e10))), introduced(avatar_definition, [new_symbols(naming, [spl144_61])])).
fof(f3270, plain, (~ (h1(e11) = h4(op1(e10, e10))) | spl144_336), inference(avatar_component_clause, [], [f3268])).
fof(f3268, plain, (spl144_336 <=> (h1(e11) = h4(op1(e10, e10)))), introduced(avatar_definition, [new_symbols(naming, [spl144_336])])).
fof(f5199, plain, (~ spl144_59 | ~ spl144_123 | ~ spl144_294 | spl144_335), inference(avatar_contradiction_clause, [], [f5198])).
fof(f5198, plain, ($false | (~ spl144_59 | ~ spl144_123 | ~ spl144_294 | spl144_335)), inference(subsumption_resolution, [], [f5197, f1723])).
fof(f1723, plain, ((e22 = op2(e20, e21)) | ~ spl144_123), inference(avatar_component_clause, [], [f1721])).
fof(f1721, plain, (spl144_123 <=> (e22 = op2(e20, e21))), introduced(avatar_definition, [new_symbols(naming, [spl144_123])])).
fof(f5197, plain, (~ (e22 = op2(e20, e21)) | (~ spl144_59 | ~ spl144_294 | spl144_335)), inference(forward_demodulation, [], [f5196, f3017])).
fof(f3017, plain, (e22 = h4(e12)), inference(forward_demodulation, [], [f1105, f1089])).
fof(f1105, plain, (op2(op2(e23, e23), e23) = h4(e12)), inference(cnf_transformation, [], [f17])).
fof(f5196, plain, (~ (op2(e20, e21) = h4(e12)) | (~ spl144_59 | ~ spl144_294 | spl144_335)), inference(forward_demodulation, [], [f5195, f1419])).
fof(f1419, plain, ((e12 = op1(e10, e11)) | ~ spl144_59), inference(avatar_component_clause, [], [f1417])).
fof(f1417, plain, (spl144_59 <=> (e12 = op1(e10, e11))), introduced(avatar_definition, [new_symbols(naming, [spl144_59])])).
fof(f5195, plain, (~ (op2(e20, e21) = h4(op1(e10, e11))) | (~ spl144_294 | spl144_335)), inference(forward_demodulation, [], [f3266, f3039])).
fof(f3266, plain, (~ (h4(op1(e10, e11)) = op2(e20, h4(e11))) | spl144_335), inference(avatar_component_clause, [], [f3264])).
fof(f3264, plain, (spl144_335 <=> (h4(op1(e10, e11)) = op2(e20, h4(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl144_335])])).
fof(f5174, plain, (~ spl144_46 | ~ spl144_110 | ~ spl144_294 | spl144_332), inference(avatar_contradiction_clause, [], [f5173])).
fof(f5173, plain, ($false | (~ spl144_46 | ~ spl144_110 | ~ spl144_294 | spl144_332)), inference(subsumption_resolution, [], [f5172, f1668])).
fof(f1668, plain, ((e21 = op2(e21, e20)) | ~ spl144_110), inference(avatar_component_clause, [], [f1666])).
fof(f1666, plain, (spl144_110 <=> (e21 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl144_110])])).
fof(f5172, plain, (~ (e21 = op2(e21, e20)) | (~ spl144_46 | ~ spl144_294 | spl144_332)), inference(forward_demodulation, [], [f5171, f3039])).
fof(f5171, plain, (~ (op2(e21, e20) = h4(e11)) | (~ spl144_46 | ~ spl144_294 | spl144_332)), inference(forward_demodulation, [], [f5170, f1364])).
fof(f1364, plain, ((e11 = op1(e11, e10)) | ~ spl144_46), inference(avatar_component_clause, [], [f1362])).
fof(f1362, plain, (spl144_46 <=> (e11 = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl144_46])])).
fof(f5170, plain, (~ (op2(e21, e20) = h4(op1(e11, e10))) | (~ spl144_294 | spl144_332)), inference(forward_demodulation, [], [f3254, f3039])).
fof(f3254, plain, (~ (h4(op1(e11, e10)) = op2(h4(e11), e20)) | spl144_332), inference(avatar_component_clause, [], [f3252])).
fof(f3252, plain, (spl144_332 <=> (h4(op1(e11, e10)) = op2(h4(e11), e20))), introduced(avatar_definition, [new_symbols(naming, [spl144_332])])).
fof(f5149, plain, (~ spl144_52 | ~ spl144_116 | spl144_333), inference(avatar_contradiction_clause, [], [f5148])).
fof(f5148, plain, ($false | (~ spl144_52 | ~ spl144_116 | spl144_333)), inference(subsumption_resolution, [], [f5147, f1693])).
fof(f1693, plain, ((e23 = op2(e20, e23)) | ~ spl144_116), inference(avatar_component_clause, [], [f1691])).
fof(f1691, plain, (spl144_116 <=> (e23 = op2(e20, e23))), introduced(avatar_definition, [new_symbols(naming, [spl144_116])])).
fof(f5147, plain, (~ (e23 = op2(e20, e23)) | (~ spl144_52 | spl144_333)), inference(forward_demodulation, [], [f5146, f1102])).
fof(f1102, plain, (e23 = h4(e13)), inference(cnf_transformation, [], [f17])).
fof(f5146, plain, (~ (op2(e20, e23) = h4(e13)) | (~ spl144_52 | spl144_333)), inference(forward_demodulation, [], [f3258, f1389])).
fof(f1389, plain, ((e13 = op1(e10, e13)) | ~ spl144_52), inference(avatar_component_clause, [], [f1387])).
fof(f1387, plain, (spl144_52 <=> (e13 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl144_52])])).
fof(f3258, plain, (~ (op2(e20, e23) = h4(op1(e10, e13))) | spl144_333), inference(avatar_component_clause, [], [f3256])).
fof(f3256, plain, (spl144_333 <=> (op2(e20, e23) = h4(op1(e10, e13)))), introduced(avatar_definition, [new_symbols(naming, [spl144_333])])).
fof(f5118, plain, (~ spl144_37 | spl144_331), inference(avatar_contradiction_clause, [], [f5117])).
fof(f5117, plain, ($false | (~ spl144_37 | spl144_331)), inference(subsumption_resolution, [], [f5116, f3028])).
fof(f5116, plain, (~ (e20 = h4(e10)) | (~ spl144_37 | spl144_331)), inference(forward_demodulation, [], [f3250, f1326])).
fof(f1326, plain, ((e10 = op1(e11, e12)) | ~ spl144_37), inference(avatar_component_clause, [], [f1324])).
fof(f1324, plain, (spl144_37 <=> (e10 = op1(e11, e12))), introduced(avatar_definition, [new_symbols(naming, [spl144_37])])).
fof(f3250, plain, (~ (e20 = h4(op1(e11, e12))) | spl144_331), inference(avatar_component_clause, [], [f3248])).
fof(f3248, plain, (spl144_331 <=> (e20 = h4(op1(e11, e12)))), introduced(avatar_definition, [new_symbols(naming, [spl144_331])])).
fof(f5087, plain, (~ spl144_44 | ~ spl144_108 | ~ spl144_294 | spl144_321), inference(avatar_contradiction_clause, [], [f5086])).
fof(f5086, plain, ($false | (~ spl144_44 | ~ spl144_108 | ~ spl144_294 | spl144_321)), inference(subsumption_resolution, [], [f5085, f1659])).
fof(f1659, plain, ((e23 = op2(e21, e21)) | ~ spl144_108), inference(avatar_component_clause, [], [f1657])).
fof(f1657, plain, (spl144_108 <=> (e23 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl144_108])])).
fof(f5085, plain, (~ (e23 = op2(e21, e21)) | (~ spl144_44 | ~ spl144_294 | spl144_321)), inference(forward_demodulation, [], [f5084, f1102])).
fof(f5084, plain, (~ (op2(e21, e21) = h4(e13)) | (~ spl144_44 | ~ spl144_294 | spl144_321)), inference(forward_demodulation, [], [f5083, f1355])).
fof(f1355, plain, ((e13 = op1(e11, e11)) | ~ spl144_44), inference(avatar_component_clause, [], [f1353])).
fof(f1353, plain, (spl144_44 <=> (e13 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl144_44])])).
fof(f5083, plain, (~ (op2(e21, e21) = h4(op1(e11, e11))) | (~ spl144_294 | spl144_321)), inference(forward_demodulation, [], [f3210, f3039])).
fof(f3210, plain, (~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | spl144_321), inference(avatar_component_clause, [], [f3208])).
fof(f3208, plain, (spl144_321 <=> (h4(op1(e11, e11)) = op2(h4(e11), h4(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl144_321])])).
fof(f5058, plain, (~ spl144_35 | spl144_330), inference(avatar_contradiction_clause, [], [f5057])).
fof(f5057, plain, ($false | (~ spl144_35 | spl144_330)), inference(subsumption_resolution, [], [f5056, f3017])).
fof(f5056, plain, (~ (e22 = h4(e12)) | (~ spl144_35 | spl144_330)), inference(forward_demodulation, [], [f3246, f1317])).
fof(f1317, plain, ((e12 = op1(e11, e13)) | ~ spl144_35), inference(avatar_component_clause, [], [f1315])).
fof(f1315, plain, (spl144_35 <=> (e12 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl144_35])])).
fof(f3246, plain, (~ (e22 = h4(op1(e11, e13))) | spl144_330), inference(avatar_component_clause, [], [f3244])).
fof(f3244, plain, (spl144_330 <=> (e22 = h4(op1(e11, e13)))), introduced(avatar_definition, [new_symbols(naming, [spl144_330])])).
fof(f5030, plain, (~ spl144_17 | ~ spl144_81 | spl144_326), inference(avatar_contradiction_clause, [], [f5029])).
fof(f5029, plain, ($false | (~ spl144_17 | ~ spl144_81 | spl144_326)), inference(subsumption_resolution, [], [f5028, f1545])).
fof(f1545, plain, ((e20 = op2(e22, e23)) | ~ spl144_81), inference(avatar_component_clause, [], [f1543])).
fof(f1543, plain, (spl144_81 <=> (e20 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl144_81])])).
fof(f5028, plain, (~ (e20 = op2(e22, e23)) | (~ spl144_17 | spl144_326)), inference(forward_demodulation, [], [f5027, f3028])).
fof(f5027, plain, (~ (op2(e22, e23) = h4(e10)) | (~ spl144_17 | spl144_326)), inference(forward_demodulation, [], [f3230, f1241])).
fof(f1241, plain, ((e10 = op1(e12, e13)) | ~ spl144_17), inference(avatar_component_clause, [], [f1239])).
fof(f1239, plain, (spl144_17 <=> (e10 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl144_17])])).
fof(f3230, plain, (~ (op2(e22, e23) = h4(op1(e12, e13))) | spl144_326), inference(avatar_component_clause, [], [f3228])).
fof(f3228, plain, (spl144_326 <=> (op2(e22, e23) = h4(op1(e12, e13)))), introduced(avatar_definition, [new_symbols(naming, [spl144_326])])).
fof(f5000, plain, (~ spl144_9 | ~ spl144_73 | ~ spl144_294 | spl144_324), inference(avatar_contradiction_clause, [], [f4999])).
fof(f4999, plain, ($false | (~ spl144_9 | ~ spl144_73 | ~ spl144_294 | spl144_324)), inference(subsumption_resolution, [], [f4998, f1511])).
fof(f1511, plain, ((e20 = op2(e23, e21)) | ~ spl144_73), inference(avatar_component_clause, [], [f1509])).
fof(f1509, plain, (spl144_73 <=> (e20 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl144_73])])).
fof(f4998, plain, (~ (e20 = op2(e23, e21)) | (~ spl144_9 | ~ spl144_294 | spl144_324)), inference(forward_demodulation, [], [f4997, f3028])).
fof(f4997, plain, (~ (op2(e23, e21) = h4(e10)) | (~ spl144_9 | ~ spl144_294 | spl144_324)), inference(forward_demodulation, [], [f4996, f1207])).
fof(f1207, plain, ((e10 = op1(e13, e11)) | ~ spl144_9), inference(avatar_component_clause, [], [f1205])).
fof(f1205, plain, (spl144_9 <=> (e10 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl144_9])])).
fof(f4996, plain, (~ (op2(e23, e21) = h4(op1(e13, e11))) | (~ spl144_294 | spl144_324)), inference(forward_demodulation, [], [f3222, f3039])).
fof(f3222, plain, (~ (h4(op1(e13, e11)) = op2(e23, h4(e11))) | spl144_324), inference(avatar_component_clause, [], [f3220])).
fof(f3220, plain, (spl144_324 <=> (h4(op1(e13, e11)) = op2(e23, h4(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl144_324])])).
fof(f4965, plain, (~ spl144_32 | ~ spl144_96 | spl144_329), inference(avatar_contradiction_clause, [], [f4964])).
fof(f4964, plain, ($false | (~ spl144_32 | ~ spl144_96 | spl144_329)), inference(subsumption_resolution, [], [f4963, f1608])).
fof(f1608, plain, ((e23 = op2(e22, e20)) | ~ spl144_96), inference(avatar_component_clause, [], [f1606])).
fof(f1606, plain, (spl144_96 <=> (e23 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl144_96])])).
fof(f4963, plain, (~ (e23 = op2(e22, e20)) | (~ spl144_32 | spl144_329)), inference(forward_demodulation, [], [f4962, f1102])).
fof(f4962, plain, (~ (op2(e22, e20) = h4(e13)) | (~ spl144_32 | spl144_329)), inference(forward_demodulation, [], [f3242, f1304])).
fof(f1304, plain, ((e13 = op1(e12, e10)) | ~ spl144_32), inference(avatar_component_clause, [], [f1302])).
fof(f1302, plain, (spl144_32 <=> (e13 = op1(e12, e10))), introduced(avatar_definition, [new_symbols(naming, [spl144_32])])).
fof(f3242, plain, (~ (op2(e22, e20) = h4(op1(e12, e10))) | spl144_329), inference(avatar_component_clause, [], [f3240])).
fof(f3240, plain, (spl144_329 <=> (op2(e22, e20) = h4(op1(e12, e10)))), introduced(avatar_definition, [new_symbols(naming, [spl144_329])])).
fof(f4932, plain, (~ spl144_23 | ~ spl144_87 | spl144_327), inference(avatar_contradiction_clause, [], [f4931])).
fof(f4931, plain, ($false | (~ spl144_23 | ~ spl144_87 | spl144_327)), inference(subsumption_resolution, [], [f4930, f4722])).
fof(f4722, plain, ((e22 = h3(e11)) | ~ spl144_87), inference(backward_demodulation, [], [f1100, f1570])).
fof(f1570, plain, ((e22 = op2(e22, e22)) | ~ spl144_87), inference(avatar_component_clause, [], [f1568])).
fof(f1568, plain, (spl144_87 <=> (e22 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl144_87])])).
fof(f1100, plain, (op2(e22, e22) = h3(e11)), inference(cnf_transformation, [], [f16])).
fof(f16, plain, ((op2(op2(e22, e22), e22) = h3(e12)) & (op2(e22, e22) = h3(e11)) & (h3(e10) = op2(op2(e22, e22), op2(op2(e22, e22), e22))) & (e22 = h3(e13))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG115+1.p', ax16)).
fof(f4930, plain, (~ (e22 = h3(e11)) | (~ spl144_23 | spl144_327)), inference(forward_demodulation, [], [f4929, f3017])).
fof(f4929, plain, (~ (h3(e11) = h4(e12)) | (~ spl144_23 | spl144_327)), inference(forward_demodulation, [], [f3234, f1266])).
fof(f1266, plain, ((e12 = op1(e12, e12)) | ~ spl144_23), inference(avatar_component_clause, [], [f1264])).
fof(f1264, plain, (spl144_23 <=> (e12 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl144_23])])).
fof(f3234, plain, (~ (h3(e11) = h4(op1(e12, e12))) | spl144_327), inference(avatar_component_clause, [], [f3232])).
fof(f3232, plain, (spl144_327 <=> (h3(e11) = h4(op1(e12, e12)))), introduced(avatar_definition, [new_symbols(naming, [spl144_327])])).
fof(f4902, plain, (~ spl144_8 | ~ spl144_72 | spl144_323), inference(avatar_contradiction_clause, [], [f4901])).
fof(f4901, plain, ($false | (~ spl144_8 | ~ spl144_72 | spl144_323)), inference(subsumption_resolution, [], [f4900, f1506])).
fof(f1506, plain, ((e23 = op2(e23, e22)) | ~ spl144_72), inference(avatar_component_clause, [], [f1504])).
fof(f1504, plain, (spl144_72 <=> (e23 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl144_72])])).
fof(f4900, plain, (~ (e23 = op2(e23, e22)) | (~ spl144_8 | spl144_323)), inference(forward_demodulation, [], [f4899, f1102])).
fof(f4899, plain, (~ (op2(e23, e22) = h4(e13)) | (~ spl144_8 | spl144_323)), inference(forward_demodulation, [], [f3218, f1202])).
fof(f1202, plain, ((e13 = op1(e13, e12)) | ~ spl144_8), inference(avatar_component_clause, [], [f1200])).
fof(f1200, plain, (spl144_8 <=> (e13 = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl144_8])])).
fof(f3218, plain, (~ (op2(e23, e22) = h4(op1(e13, e12))) | spl144_323), inference(avatar_component_clause, [], [f3216])).
fof(f3216, plain, (spl144_323 <=> (op2(e23, e22) = h4(op1(e13, e12)))), introduced(avatar_definition, [new_symbols(naming, [spl144_323])])).
fof(f4870, plain, (~ spl144_42 | ~ spl144_46), inference(avatar_split_clause, [], [f4867, f1362, f1345])).
fof(f1345, plain, (spl144_42 <=> (e11 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl144_42])])).
fof(f4867, plain, (~ (e11 = op1(e11, e11)) | ~ spl144_46), inference(backward_demodulation, [], [f438, f1364])).
fof(f438, plain, ~ (op1(e11, e10) = op1(e11, e11)), inference(cnf_transformation, [], [f5])).
fof(f5, plain, (~ (op1(e13, e12) = op1(e13, e13)) & ~ (op1(e13, e11) = op1(e13, e13)) & ~ (op1(e13, e11) = op1(e13, e12)) & ~ (op1(e13, e10) = op1(e13, e13)) & ~ (op1(e13, e10) = op1(e13, e12)) & ~ (op1(e13, e10) = op1(e13, e11)) & ~ (op1(e12, e12) = op1(e12, e13)) & ~ (op1(e12, e11) = op1(e12, e13)) & ~ (op1(e12, e11) = op1(e12, e12)) & ~ (op1(e12, e10) = op1(e12, e13)) & ~ (op1(e12, e10) = op1(e12, e12)) & ~ (op1(e12, e10) = op1(e12, e11)) & ~ (op1(e11, e12) = op1(e11, e13)) & ~ (op1(e11, e11) = op1(e11, e13)) & ~ (op1(e11, e11) = op1(e11, e12)) & ~ (op1(e11, e10) = op1(e11, e13)) & ~ (op1(e11, e10) = op1(e11, e12)) & ~ (op1(e11, e10) = op1(e11, e11)) & ~ (op1(e10, e12) = op1(e10, e13)) & ~ (op1(e10, e11) = op1(e10, e13)) & ~ (op1(e10, e11) = op1(e10, e12)) & ~ (op1(e10, e10) = op1(e10, e13)) & ~ (op1(e10, e10) = op1(e10, e12)) & ~ (op1(e10, e10) = op1(e10, e11)) & ~ (op1(e12, e13) = op1(e13, e13)) & ~ (op1(e11, e13) = op1(e13, e13)) & ~ (op1(e11, e13) = op1(e12, e13)) & ~ (op1(e10, e13) = op1(e13, e13)) & ~ (op1(e10, e13) = op1(e12, e13)) & ~ (op1(e10, e13) = op1(e11, e13)) & ~ (op1(e12, e12) = op1(e13, e12)) & ~ (op1(e11, e12) = op1(e13, e12)) & ~ (op1(e11, e12) = op1(e12, e12)) & ~ (op1(e10, e12) = op1(e13, e12)) & ~ (op1(e10, e12) = op1(e12, e12)) & ~ (op1(e10, e12) = op1(e11, e12)) & ~ (op1(e12, e11) = op1(e13, e11)) & ~ (op1(e11, e11) = op1(e13, e11)) & ~ (op1(e11, e11) = op1(e12, e11)) & ~ (op1(e10, e11) = op1(e13, e11)) & ~ (op1(e10, e11) = op1(e12, e11)) & ~ (op1(e10, e11) = op1(e11, e11)) & ~ (op1(e12, e10) = op1(e13, e10)) & ~ (op1(e11, e10) = op1(e13, e10)) & ~ (op1(e11, e10) = op1(e12, e10)) & ~ (op1(e10, e10) = op1(e13, e10)) & ~ (op1(e10, e10) = op1(e12, e10)) & ~ (op1(e10, e10) = op1(e11, e10))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG115+1.p', ax5)).
fof(f4858, plain, (~ spl144_2 | spl144_322), inference(avatar_contradiction_clause, [], [f4857])).
fof(f4857, plain, ($false | (~ spl144_2 | spl144_322)), inference(trivial_inequality_removal, [], [f4856])).
fof(f4856, plain, (~ (h4(e11) = h4(e11)) | (~ spl144_2 | spl144_322)), inference(forward_demodulation, [], [f3214, f1177])).
fof(f1177, plain, ((e11 = op1(e13, e13)) | ~ spl144_2), inference(avatar_component_clause, [], [f1175])).
fof(f1175, plain, (spl144_2 <=> (e11 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl144_2])])).
fof(f3214, plain, (~ (h4(e11) = h4(op1(e13, e13))) | spl144_322), inference(avatar_component_clause, [], [f3212])).
fof(f3212, plain, (spl144_322 <=> (h4(e11) = h4(op1(e13, e13)))), introduced(avatar_definition, [new_symbols(naming, [spl144_322])])).
fof(f4829, plain, (~ spl144_44 | ~ spl144_48), inference(avatar_split_clause, [], [f4827, f1370, f1353])).
fof(f1370, plain, (spl144_48 <=> (e13 = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl144_48])])).
fof(f4827, plain, (~ (e13 = op1(e11, e11)) | ~ spl144_48), inference(backward_demodulation, [], [f438, f1372])).
fof(f1372, plain, ((e13 = op1(e11, e10)) | ~ spl144_48), inference(avatar_component_clause, [], [f1370])).
fof(f4825, plain, (~ spl144_61 | ~ spl144_64), inference(avatar_contradiction_clause, [], [f4824])).
fof(f4824, plain, ($false | (~ spl144_61 | ~ spl144_64)), inference(subsumption_resolution, [], [f4823, f506])).
fof(f506, plain, ~ (e10 = e13), inference(cnf_transformation, [], [f7])).
fof(f7, plain, (~ (e12 = e13) & ~ (e11 = e13) & ~ (e11 = e12) & ~ (e10 = e13) & ~ (e10 = e12) & ~ (e10 = e11)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG115+1.p', ax7)).
fof(f4823, plain, ((e10 = e13) | (~ spl144_61 | ~ spl144_64)), inference(forward_demodulation, [], [f1440, f1428])).
fof(f1440, plain, ((op1(e10, e10) = e13) | ~ spl144_64), inference(avatar_component_clause, [], [f1438])).
fof(f1438, plain, (spl144_64 <=> (op1(e10, e10) = e13)), introduced(avatar_definition, [new_symbols(naming, [spl144_64])])).
fof(f4808, plain, (~ spl144_26 | ~ spl144_90 | ~ spl144_294 | spl144_328), inference(avatar_contradiction_clause, [], [f4807])).
fof(f4807, plain, ($false | (~ spl144_26 | ~ spl144_90 | ~ spl144_294 | spl144_328)), inference(subsumption_resolution, [], [f4806, f1583])).
fof(f1583, plain, ((e21 = op2(e22, e21)) | ~ spl144_90), inference(avatar_component_clause, [], [f1581])).
fof(f1581, plain, (spl144_90 <=> (e21 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl144_90])])).
fof(f4806, plain, (~ (e21 = op2(e22, e21)) | (~ spl144_26 | ~ spl144_294 | spl144_328)), inference(forward_demodulation, [], [f4805, f3039])).
fof(f4805, plain, (~ (op2(e22, e21) = h4(e11)) | (~ spl144_26 | ~ spl144_294 | spl144_328)), inference(forward_demodulation, [], [f4804, f1279])).
fof(f1279, plain, ((e11 = op1(e12, e11)) | ~ spl144_26), inference(avatar_component_clause, [], [f1277])).
fof(f1277, plain, (spl144_26 <=> (e11 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl144_26])])).
fof(f4804, plain, (~ (op2(e22, e21) = h4(op1(e12, e11))) | (~ spl144_294 | spl144_328)), inference(forward_demodulation, [], [f3238, f3039])).
fof(f3238, plain, (~ (h4(op1(e12, e11)) = op2(e22, h4(e11))) | spl144_328), inference(avatar_component_clause, [], [f3236])).
fof(f3236, plain, (spl144_328 <=> (h4(op1(e12, e11)) = op2(e22, h4(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl144_328])])).
fof(f4774, plain, (~ spl144_7 | ~ spl144_23), inference(avatar_split_clause, [], [f4773, f1264, f1196])).
fof(f1196, plain, (spl144_7 <=> (e12 = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl144_7])])).
fof(f4773, plain, (~ (e12 = op1(e13, e12)) | ~ spl144_23), inference(backward_demodulation, [], [f425, f1266])).
fof(f425, plain, ~ (op1(e12, e12) = op1(e13, e12)), inference(cnf_transformation, [], [f5])).
fof(f4718, plain, (~ spl144_95 | ~ spl144_96), inference(avatar_contradiction_clause, [], [f4717])).
fof(f4717, plain, ($false | (~ spl144_95 | ~ spl144_96)), inference(subsumption_resolution, [], [f4716, f515])).
fof(f515, plain, ~ (e22 = e23), inference(cnf_transformation, [], [f8])).
fof(f8, plain, (~ (e22 = e23) & ~ (e21 = e23) & ~ (e21 = e22) & ~ (e20 = e23) & ~ (e20 = e22) & ~ (e20 = e21)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG115+1.p', ax8)).
fof(f4716, plain, ((e22 = e23) | (~ spl144_95 | ~ spl144_96)), inference(backward_demodulation, [], [f1608, f1604])).
fof(f1604, plain, ((e22 = op2(e22, e20)) | ~ spl144_95), inference(avatar_component_clause, [], [f1602])).
fof(f1602, plain, (spl144_95 <=> (e22 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl144_95])])).
fof(f4708, plain, (~ spl144_94 | ~ spl144_110), inference(avatar_split_clause, [], [f4704, f1666, f1598])).
fof(f1598, plain, (spl144_94 <=> (e21 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl144_94])])).
fof(f4704, plain, (~ (e21 = op2(e22, e20)) | ~ spl144_110), inference(backward_demodulation, [], [f459, f1668])).
fof(f459, plain, ~ (op2(e21, e20) = op2(e22, e20)), inference(cnf_transformation, [], [f6])).
fof(f6, plain, (~ (op2(e23, e22) = op2(e23, e23)) & ~ (op2(e23, e21) = op2(e23, e23)) & ~ (op2(e23, e21) = op2(e23, e22)) & ~ (op2(e23, e20) = op2(e23, e23)) & ~ (op2(e23, e20) = op2(e23, e22)) & ~ (op2(e23, e20) = op2(e23, e21)) & ~ (op2(e22, e22) = op2(e22, e23)) & ~ (op2(e22, e21) = op2(e22, e23)) & ~ (op2(e22, e21) = op2(e22, e22)) & ~ (op2(e22, e20) = op2(e22, e23)) & ~ (op2(e22, e20) = op2(e22, e22)) & ~ (op2(e22, e20) = op2(e22, e21)) & ~ (op2(e21, e22) = op2(e21, e23)) & ~ (op2(e21, e21) = op2(e21, e23)) & ~ (op2(e21, e21) = op2(e21, e22)) & ~ (op2(e21, e20) = op2(e21, e23)) & ~ (op2(e21, e20) = op2(e21, e22)) & ~ (op2(e21, e20) = op2(e21, e21)) & ~ (op2(e20, e22) = op2(e20, e23)) & ~ (op2(e20, e21) = op2(e20, e23)) & ~ (op2(e20, e21) = op2(e20, e22)) & ~ (op2(e20, e20) = op2(e20, e23)) & ~ (op2(e20, e20) = op2(e20, e22)) & ~ (op2(e20, e20) = op2(e20, e21)) & ~ (op2(e22, e23) = op2(e23, e23)) & ~ (op2(e21, e23) = op2(e23, e23)) & ~ (op2(e21, e23) = op2(e22, e23)) & ~ (op2(e20, e23) = op2(e23, e23)) & ~ (op2(e20, e23) = op2(e22, e23)) & ~ (op2(e20, e23) = op2(e21, e23)) & ~ (op2(e22, e22) = op2(e23, e22)) & ~ (op2(e21, e22) = op2(e23, e22)) & ~ (op2(e21, e22) = op2(e22, e22)) & ~ (op2(e20, e22) = op2(e23, e22)) & ~ (op2(e20, e22) = op2(e22, e22)) & ~ (op2(e20, e22) = op2(e21, e22)) & ~ (op2(e22, e21) = op2(e23, e21)) & ~ (op2(e21, e21) = op2(e23, e21)) & ~ (op2(e21, e21) = op2(e22, e21)) & ~ (op2(e20, e21) = op2(e23, e21)) & ~ (op2(e20, e21) = op2(e22, e21)) & ~ (op2(e20, e21) = op2(e21, e21)) & ~ (op2(e22, e20) = op2(e23, e20)) & ~ (op2(e21, e20) = op2(e23, e20)) & ~ (op2(e21, e20) = op2(e22, e20)) & ~ (op2(e20, e20) = op2(e23, e20)) & ~ (op2(e20, e20) = op2(e22, e20)) & ~ (op2(e20, e20) = op2(e21, e20))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG115+1.p', ax6)).
fof(f4700, plain, (~ spl144_118 | ~ spl144_120), inference(avatar_contradiction_clause, [], [f4699])).
fof(f4699, plain, ($false | (~ spl144_118 | ~ spl144_120)), inference(subsumption_resolution, [], [f4698, f514])).
fof(f514, plain, ~ (e21 = e23), inference(cnf_transformation, [], [f8])).
fof(f4698, plain, ((e21 = e23) | (~ spl144_118 | ~ spl144_120)), inference(forward_demodulation, [], [f1710, f1702])).
fof(f1710, plain, ((e23 = op2(e20, e22)) | ~ spl144_120), inference(avatar_component_clause, [], [f1708])).
fof(f1708, plain, (spl144_120 <=> (e23 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl144_120])])).
fof(f4687, plain, (~ spl144_112 | ~ spl144_374), inference(avatar_split_clause, [], [f4680, f3516, f1674])).
fof(f1674, plain, (spl144_112 <=> (e23 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl144_112])])).
fof(f3516, plain, (spl144_374 <=> (e23 = h2(e11))), introduced(avatar_definition, [new_symbols(naming, [spl144_374])])).
fof(f4680, plain, (~ (e23 = op2(e21, e20)) | ~ spl144_374), inference(backward_demodulation, [], [f3002, f3517])).
fof(f3517, plain, ((e23 = h2(e11)) | ~ spl144_374), inference(avatar_component_clause, [], [f3516])).
fof(f3002, plain, ~ (op2(e21, e20) = h2(e11)), inference(backward_demodulation, [], [f486, f1096])).
fof(f1096, plain, (op2(e21, e21) = h2(e11)), inference(cnf_transformation, [], [f15])).
fof(f15, plain, ((op2(op2(e21, e21), e21) = h2(e12)) & (op2(e21, e21) = h2(e11)) & (h2(e10) = op2(op2(e21, e21), op2(op2(e21, e21), e21))) & (e21 = h2(e13))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG115+1.p', ax15)).
fof(f486, plain, ~ (op2(e21, e20) = op2(e21, e21)), inference(cnf_transformation, [], [f6])).
fof(f4686, plain, (~ spl144_92 | ~ spl144_374), inference(avatar_split_clause, [], [f4679, f3516, f1589])).
fof(f1589, plain, (spl144_92 <=> (e23 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl144_92])])).
fof(f4679, plain, (~ (e23 = op2(e22, e21)) | ~ spl144_374), inference(backward_demodulation, [], [f3000, f3517])).
fof(f3000, plain, ~ (op2(e22, e21) = h2(e11)), inference(backward_demodulation, [], [f465, f1096])).
fof(f465, plain, ~ (op2(e21, e21) = op2(e22, e21)), inference(cnf_transformation, [], [f6])).
fof(f4685, plain, (~ spl144_124 | ~ spl144_374), inference(avatar_split_clause, [], [f4678, f3516, f1725])).
fof(f1725, plain, (spl144_124 <=> (e23 = op2(e20, e21))), introduced(avatar_definition, [new_symbols(naming, [spl144_124])])).
fof(f4678, plain, (~ (e23 = op2(e20, e21)) | ~ spl144_374), inference(backward_demodulation, [], [f2999, f3517])).
fof(f2999, plain, ~ (op2(e20, e21) = h2(e11)), inference(backward_demodulation, [], [f462, f1096])).
fof(f462, plain, ~ (op2(e20, e21) = op2(e21, e21)), inference(cnf_transformation, [], [f6])).
fof(f4643, plain, (~ spl144_61 | ~ spl144_63), inference(avatar_contradiction_clause, [], [f4642])).
fof(f4642, plain, ($false | (~ spl144_61 | ~ spl144_63)), inference(subsumption_resolution, [], [f4641, f505])).
fof(f505, plain, ~ (e10 = e12), inference(cnf_transformation, [], [f7])).
fof(f4641, plain, ((e10 = e12) | (~ spl144_61 | ~ spl144_63)), inference(forward_demodulation, [], [f1436, f1428])).
fof(f1436, plain, ((op1(e10, e10) = e12) | ~ spl144_63), inference(avatar_component_clause, [], [f1434])).
fof(f1434, plain, (spl144_63 <=> (op1(e10, e10) = e12)), introduced(avatar_definition, [new_symbols(naming, [spl144_63])])).
fof(f4596, plain, (~ spl144_125 | ~ spl144_127), inference(avatar_contradiction_clause, [], [f4595])).
fof(f4595, plain, ($false | (~ spl144_125 | ~ spl144_127)), inference(subsumption_resolution, [], [f4594, f511])).
fof(f511, plain, ~ (e20 = e22), inference(cnf_transformation, [], [f8])).
fof(f4594, plain, ((e20 = e22) | (~ spl144_125 | ~ spl144_127)), inference(forward_demodulation, [], [f1740, f1732])).
fof(f1740, plain, ((op2(e20, e20) = e22) | ~ spl144_127), inference(avatar_component_clause, [], [f1738])).
fof(f1738, plain, (spl144_127 <=> (op2(e20, e20) = e22)), introduced(avatar_definition, [new_symbols(naming, [spl144_127])])).
fof(f4593, plain, (~ spl144_125 | ~ spl144_128), inference(avatar_contradiction_clause, [], [f4592])).
fof(f4592, plain, ($false | (~ spl144_125 | ~ spl144_128)), inference(subsumption_resolution, [], [f4591, f512])).
fof(f512, plain, ~ (e20 = e23), inference(cnf_transformation, [], [f8])).
fof(f4591, plain, ((e20 = e23) | (~ spl144_125 | ~ spl144_128)), inference(forward_demodulation, [], [f1744, f1732])).
fof(f1744, plain, ((op2(e20, e20) = e23) | ~ spl144_128), inference(avatar_component_clause, [], [f1742])).
fof(f1742, plain, (spl144_128 <=> (op2(e20, e20) = e23)), introduced(avatar_definition, [new_symbols(naming, [spl144_128])])).
fof(f4571, plain, (spl144_99 | ~ spl144_294), inference(avatar_split_clause, [], [f4301, f3038, f1619])).
fof(f1619, plain, (spl144_99 <=> (e22 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl144_99])])).
fof(f4301, plain, ((e22 = op2(e21, e23)) | ~ spl144_294), inference(backward_demodulation, [], [f3024, f3039])).
fof(f4510, plain, (~ spl144_22 | ~ spl144_23), inference(avatar_contradiction_clause, [], [f4509])).
fof(f4509, plain, ($false | (~ spl144_22 | ~ spl144_23)), inference(subsumption_resolution, [], [f4508, f507])).
fof(f507, plain, ~ (e11 = e12), inference(cnf_transformation, [], [f7])).
fof(f4508, plain, ((e11 = e12) | (~ spl144_22 | ~ spl144_23)), inference(backward_demodulation, [], [f1266, f1262])).
fof(f1262, plain, ((e11 = op1(e12, e12)) | ~ spl144_22), inference(avatar_component_clause, [], [f1260])).
fof(f1260, plain, (spl144_22 <=> (e11 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl144_22])])).
fof(f4507, plain, (~ spl144_23 | ~ spl144_24), inference(avatar_contradiction_clause, [], [f4506])).
fof(f4506, plain, ($false | (~ spl144_23 | ~ spl144_24)), inference(subsumption_resolution, [], [f4505, f509])).
fof(f509, plain, ~ (e12 = e13), inference(cnf_transformation, [], [f7])).
fof(f4505, plain, ((e12 = e13) | (~ spl144_23 | ~ spl144_24)), inference(forward_demodulation, [], [f1270, f1266])).
fof(f1270, plain, ((e13 = op1(e12, e12)) | ~ spl144_24), inference(avatar_component_clause, [], [f1268])).
fof(f1268, plain, (spl144_24 <=> (e13 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl144_24])])).
fof(f4485, plain, (~ spl144_66 | ~ spl144_68), inference(avatar_contradiction_clause, [], [f4484])).
fof(f4484, plain, ($false | (~ spl144_66 | ~ spl144_68)), inference(subsumption_resolution, [], [f4483, f514])).
fof(f4483, plain, ((e21 = e23) | (~ spl144_66 | ~ spl144_68)), inference(forward_demodulation, [], [f1489, f1481])).
fof(f1481, plain, ((e21 = op2(e23, e23)) | ~ spl144_66), inference(avatar_component_clause, [], [f1479])).
fof(f1479, plain, (spl144_66 <=> (e21 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl144_66])])).
fof(f1489, plain, ((e23 = op2(e23, e23)) | ~ spl144_68), inference(avatar_component_clause, [], [f1487])).
fof(f1487, plain, (spl144_68 <=> (e23 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl144_68])])).
fof(f4469, plain, (~ spl144_101 | ~ spl144_102), inference(avatar_contradiction_clause, [], [f4468])).
fof(f4468, plain, ($false | (~ spl144_101 | ~ spl144_102)), inference(subsumption_resolution, [], [f4467, f510])).
fof(f510, plain, ~ (e20 = e21), inference(cnf_transformation, [], [f8])).
fof(f4467, plain, ((e20 = e21) | (~ spl144_101 | ~ spl144_102)), inference(backward_demodulation, [], [f1634, f1630])).
fof(f1630, plain, ((e20 = op2(e21, e22)) | ~ spl144_101), inference(avatar_component_clause, [], [f1628])).
fof(f1628, plain, (spl144_101 <=> (e20 = op2(e21, e22))), introduced(avatar_definition, [new_symbols(naming, [spl144_101])])).
fof(f1634, plain, ((e21 = op2(e21, e22)) | ~ spl144_102), inference(avatar_component_clause, [], [f1632])).
fof(f1632, plain, (spl144_102 <=> (e21 = op2(e21, e22))), introduced(avatar_definition, [new_symbols(naming, [spl144_102])])).
fof(f4461, plain, (~ spl144_107 | spl144_307), inference(avatar_contradiction_clause, [], [f4460])).
fof(f4460, plain, ($false | (~ spl144_107 | spl144_307)), inference(subsumption_resolution, [], [f4459, f3110])).
fof(f3110, plain, (~ (e22 = h2(e11)) | spl144_307), inference(avatar_component_clause, [], [f3108])).
fof(f3108, plain, (spl144_307 <=> (e22 = h2(e11))), introduced(avatar_definition, [new_symbols(naming, [spl144_307])])).
fof(f4459, plain, ((e22 = h2(e11)) | ~ spl144_107), inference(backward_demodulation, [], [f1096, f1655])).
fof(f1655, plain, ((e22 = op2(e21, e21)) | ~ spl144_107), inference(avatar_component_clause, [], [f1653])).
fof(f1653, plain, (spl144_107 <=> (e22 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl144_107])])).
fof(f4417, plain, (~ spl144_70 | ~ spl144_294), inference(avatar_split_clause, [], [f4416, f3038, f1496])).
fof(f1496, plain, (spl144_70 <=> (e21 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl144_70])])).
fof(f4416, plain, (~ (e21 = op2(e23, e22)) | ~ spl144_294), inference(forward_demodulation, [], [f3023, f3039])).
fof(f3023, plain, ~ (op2(e23, e22) = h4(e11)), inference(backward_demodulation, [], [f503, f1104])).
fof(f503, plain, ~ (op2(e23, e22) = op2(e23, e23)), inference(cnf_transformation, [], [f6])).
fof(f4411, plain, (~ spl144_23 | ~ spl144_55), inference(avatar_contradiction_clause, [], [f4410])).
fof(f4410, plain, ($false | (~ spl144_23 | ~ spl144_55)), inference(subsumption_resolution, [], [f4409, f1402])).
fof(f1402, plain, ((e12 = op1(e10, e12)) | ~ spl144_55), inference(avatar_component_clause, [], [f1400])).
fof(f1400, plain, (spl144_55 <=> (e12 = op1(e10, e12))), introduced(avatar_definition, [new_symbols(naming, [spl144_55])])).
fof(f4409, plain, (~ (e12 = op1(e10, e12)) | ~ spl144_23), inference(forward_demodulation, [], [f421, f1266])).
fof(f421, plain, ~ (op1(e10, e12) = op1(e12, e12)), inference(cnf_transformation, [], [f5])).
fof(f4394, plain, (~ spl144_23 | ~ spl144_31), inference(avatar_contradiction_clause, [], [f4393])).
fof(f4393, plain, ($false | (~ spl144_23 | ~ spl144_31)), inference(subsumption_resolution, [], [f4392, f1300])).
fof(f1300, plain, ((e12 = op1(e12, e10)) | ~ spl144_31), inference(avatar_component_clause, [], [f1298])).
fof(f1298, plain, (spl144_31 <=> (e12 = op1(e12, e10))), introduced(avatar_definition, [new_symbols(naming, [spl144_31])])).
fof(f4392, plain, (~ (e12 = op1(e12, e10)) | ~ spl144_23), inference(forward_demodulation, [], [f445, f1266])).
fof(f445, plain, ~ (op1(e12, e10) = op1(e12, e12)), inference(cnf_transformation, [], [f5])).
fof(f4383, plain, (~ spl144_6 | ~ spl144_2), inference(avatar_split_clause, [], [f4382, f1175, f1192])).
fof(f1192, plain, (spl144_6 <=> (e11 = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl144_6])])).
fof(f4382, plain, (~ (e11 = op1(e13, e12)) | ~ spl144_2), inference(forward_demodulation, [], [f455, f1177])).
fof(f455, plain, ~ (op1(e13, e12) = op1(e13, e13)), inference(cnf_transformation, [], [f5])).
fof(f4380, plain, (~ spl144_1 | ~ spl144_2), inference(avatar_contradiction_clause, [], [f4379])).
fof(f4379, plain, ($false | (~ spl144_1 | ~ spl144_2)), inference(subsumption_resolution, [], [f4378, f504])).
fof(f504, plain, ~ (e10 = e11), inference(cnf_transformation, [], [f7])).
fof(f4378, plain, ((e10 = e11) | (~ spl144_1 | ~ spl144_2)), inference(backward_demodulation, [], [f1177, f1173])).
fof(f1173, plain, ((e10 = op1(e13, e13)) | ~ spl144_1), inference(avatar_component_clause, [], [f1171])).
fof(f1171, plain, (spl144_1 <=> (e10 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl144_1])])).
fof(f4358, plain, (~ spl144_37 | ~ spl144_38), inference(avatar_contradiction_clause, [], [f4357])).
fof(f4357, plain, ($false | (~ spl144_37 | ~ spl144_38)), inference(subsumption_resolution, [], [f4356, f504])).
fof(f4356, plain, ((e10 = e11) | (~ spl144_37 | ~ spl144_38)), inference(backward_demodulation, [], [f1330, f1326])).
fof(f1330, plain, ((e11 = op1(e11, e12)) | ~ spl144_38), inference(avatar_component_clause, [], [f1328])).
fof(f1328, plain, (spl144_38 <=> (e11 = op1(e11, e12))), introduced(avatar_definition, [new_symbols(naming, [spl144_38])])).
fof(f4329, plain, (~ spl144_86 | ~ spl144_87), inference(avatar_contradiction_clause, [], [f4328])).
fof(f4328, plain, ($false | (~ spl144_86 | ~ spl144_87)), inference(subsumption_resolution, [], [f4327, f513])).
fof(f513, plain, ~ (e21 = e22), inference(cnf_transformation, [], [f8])).
fof(f4327, plain, ((e21 = e22) | (~ spl144_86 | ~ spl144_87)), inference(backward_demodulation, [], [f1570, f1566])).
fof(f1566, plain, ((e21 = op2(e22, e22)) | ~ spl144_86), inference(avatar_component_clause, [], [f1564])).
fof(f1564, plain, (spl144_86 <=> (e21 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl144_86])])).
fof(f4319, plain, (spl144_374 | ~ spl144_108), inference(avatar_split_clause, [], [f4318, f1657, f3516])).
fof(f4318, plain, ((e23 = h2(e11)) | ~ spl144_108), inference(backward_demodulation, [], [f1096, f1659])).
fof(f4306, plain, (spl144_101 | ~ spl144_294), inference(avatar_split_clause, [], [f4302, f3038, f1628])).
fof(f4302, plain, ((e20 = op2(e21, e22)) | ~ spl144_294), inference(backward_demodulation, [], [f3025, f3039])).
fof(f4296, plain, (spl144_294 | ~ spl144_66), inference(avatar_split_clause, [], [f4047, f1479, f3038])).
fof(f4047, plain, ((e21 = h4(e11)) | ~ spl144_66), inference(forward_demodulation, [], [f1104, f1481])).
fof(f4267, plain, (~ spl144_91 | ~ spl144_87), inference(avatar_split_clause, [], [f4266, f1568, f1585])).
fof(f1585, plain, (spl144_91 <=> (e22 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl144_91])])).
fof(f4266, plain, (~ (e22 = op2(e22, e21)) | ~ spl144_87), inference(forward_demodulation, [], [f3012, f4257])).
fof(f4257, plain, ((e22 = h3(e11)) | ~ spl144_87), inference(forward_demodulation, [], [f1100, f1570])).
fof(f3012, plain, ~ (op2(e22, e21) = h3(e11)), inference(backward_demodulation, [], [f495, f1100])).
fof(f495, plain, ~ (op2(e22, e21) = op2(e22, e22)), inference(cnf_transformation, [], [f6])).
fof(f4227, plain, (~ spl144_51 | ~ spl144_35), inference(avatar_split_clause, [], [f4034, f1315, f1383])).
fof(f1383, plain, (spl144_51 <=> (e12 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl144_51])])).
fof(f4034, plain, (~ (e12 = op1(e10, e13)) | ~ spl144_35), inference(forward_demodulation, [], [f426, f1317])).
fof(f426, plain, ~ (op1(e10, e13) = op1(e11, e13)), inference(cnf_transformation, [], [f5])).
fof(f4214, plain, (~ spl144_43 | ~ spl144_35), inference(avatar_split_clause, [], [f4213, f1315, f1349])).
fof(f1349, plain, (spl144_43 <=> (e12 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl144_43])])).
fof(f4213, plain, (~ (e12 = op1(e11, e11)) | ~ spl144_35), inference(forward_demodulation, [], [f442, f1317])).
fof(f442, plain, ~ (op1(e11, e11) = op1(e11, e13)), inference(cnf_transformation, [], [f5])).
fof(f4201, plain, (~ spl144_20 | ~ spl144_32), inference(avatar_split_clause, [], [f4200, f1302, f1251])).
fof(f1251, plain, (spl144_20 <=> (e13 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl144_20])])).
fof(f4200, plain, (~ (e13 = op1(e12, e13)) | ~ spl144_32), inference(forward_demodulation, [], [f446, f1304])).
fof(f446, plain, ~ (op1(e12, e10) = op1(e12, e13)), inference(cnf_transformation, [], [f5])).
fof(f4156, plain, (~ spl144_54 | ~ spl144_58), inference(avatar_split_clause, [], [f4154, f1413, f1396])).
fof(f1413, plain, (spl144_58 <=> (e11 = op1(e10, e11))), introduced(avatar_definition, [new_symbols(naming, [spl144_58])])).
fof(f4154, plain, (~ (e11 = op1(e10, e12)) | ~ spl144_58), inference(backward_demodulation, [], [f435, f1415])).
fof(f1415, plain, ((e11 = op1(e10, e11)) | ~ spl144_58), inference(avatar_component_clause, [], [f1413])).
fof(f435, plain, ~ (op1(e10, e11) = op1(e10, e12)), inference(cnf_transformation, [], [f5])).
fof(f4143, plain, (~ spl144_81 | ~ spl144_82), inference(avatar_contradiction_clause, [], [f4142])).
fof(f4142, plain, ($false | (~ spl144_81 | ~ spl144_82)), inference(subsumption_resolution, [], [f4141, f510])).
fof(f4141, plain, ((e20 = e21) | (~ spl144_81 | ~ spl144_82)), inference(backward_demodulation, [], [f1549, f1545])).
fof(f1549, plain, ((e21 = op2(e22, e23)) | ~ spl144_82), inference(avatar_component_clause, [], [f1547])).
fof(f1547, plain, (spl144_82 <=> (e21 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl144_82])])).
fof(f4139, plain, (~ spl144_87 | ~ spl144_88), inference(avatar_contradiction_clause, [], [f4138])).
fof(f4138, plain, ($false | (~ spl144_87 | ~ spl144_88)), inference(subsumption_resolution, [], [f4137, f515])).
fof(f4137, plain, ((e22 = e23) | (~ spl144_87 | ~ spl144_88)), inference(backward_demodulation, [], [f1574, f1570])).
fof(f1574, plain, ((e23 = op2(e22, e22)) | ~ spl144_88), inference(avatar_component_clause, [], [f1572])).
fof(f1572, plain, (spl144_88 <=> (e23 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl144_88])])).
fof(f4126, plain, (~ spl144_101 | ~ spl144_104), inference(avatar_contradiction_clause, [], [f4125])).
fof(f4125, plain, ($false | (~ spl144_101 | ~ spl144_104)), inference(subsumption_resolution, [], [f4124, f512])).
fof(f4124, plain, ((e20 = e23) | (~ spl144_101 | ~ spl144_104)), inference(forward_demodulation, [], [f1642, f1630])).
fof(f1642, plain, ((e23 = op2(e21, e22)) | ~ spl144_104), inference(avatar_component_clause, [], [f1640])).
fof(f1640, plain, (spl144_104 <=> (e23 = op2(e21, e22))), introduced(avatar_definition, [new_symbols(naming, [spl144_104])])).
fof(f4092, plain, (~ spl144_109 | ~ spl144_101), inference(avatar_split_clause, [], [f4091, f1628, f1662])).
fof(f1662, plain, (spl144_109 <=> (e20 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl144_109])])).
fof(f4091, plain, (~ (e20 = op2(e21, e20)) | ~ spl144_101), inference(forward_demodulation, [], [f487, f1630])).
fof(f487, plain, ~ (op2(e21, e20) = op2(e21, e22)), inference(cnf_transformation, [], [f6])).
fof(f4090, plain, (~ spl144_111 | ~ spl144_99), inference(avatar_split_clause, [], [f4089, f1619, f1670])).
fof(f1670, plain, (spl144_111 <=> (e22 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl144_111])])).
fof(f4089, plain, (~ (e22 = op2(e21, e20)) | ~ spl144_99), inference(forward_demodulation, [], [f488, f1621])).
fof(f1621, plain, ((e22 = op2(e21, e23)) | ~ spl144_99), inference(avatar_component_clause, [], [f1619])).
fof(f488, plain, ~ (op2(e21, e20) = op2(e21, e23)), inference(cnf_transformation, [], [f6])).
fof(f4087, plain, (~ spl144_307 | ~ spl144_99), inference(avatar_split_clause, [], [f4086, f1619, f3108])).
fof(f4086, plain, (~ (e22 = h2(e11)) | ~ spl144_99), inference(forward_demodulation, [], [f3004, f1621])).
fof(f3004, plain, ~ (op2(e21, e23) = h2(e11)), inference(backward_demodulation, [], [f490, f1096])).
fof(f490, plain, ~ (op2(e21, e21) = op2(e21, e23)), inference(cnf_transformation, [], [f6])).
fof(f4072, plain, (~ spl144_75 | ~ spl144_79), inference(avatar_contradiction_clause, [], [f4071])).
fof(f4071, plain, ($false | (~ spl144_75 | ~ spl144_79)), inference(subsumption_resolution, [], [f4070, f1536])).
fof(f1536, plain, ((e22 = op2(e23, e20)) | ~ spl144_79), inference(avatar_component_clause, [], [f1534])).
fof(f1534, plain, (spl144_79 <=> (e22 = op2(e23, e20))), introduced(avatar_definition, [new_symbols(naming, [spl144_79])])).
fof(f4070, plain, (~ (e22 = op2(e23, e20)) | ~ spl144_75), inference(forward_demodulation, [], [f498, f1519])).
fof(f1519, plain, ((e22 = op2(e23, e21)) | ~ spl144_75), inference(avatar_component_clause, [], [f1517])).
fof(f1517, plain, (spl144_75 <=> (e22 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl144_75])])).
fof(f498, plain, ~ (op2(e23, e20) = op2(e23, e21)), inference(cnf_transformation, [], [f6])).
fof(f4064, plain, (~ spl144_15 | ~ spl144_79 | spl144_325), inference(avatar_contradiction_clause, [], [f4063])).
fof(f4063, plain, ($false | (~ spl144_15 | ~ spl144_79 | spl144_325)), inference(subsumption_resolution, [], [f4062, f1536])).
fof(f4062, plain, (~ (e22 = op2(e23, e20)) | (~ spl144_15 | spl144_325)), inference(forward_demodulation, [], [f4061, f3017])).
fof(f4061, plain, (~ (op2(e23, e20) = h4(e12)) | (~ spl144_15 | spl144_325)), inference(forward_demodulation, [], [f3226, f1232])).
fof(f1232, plain, ((e12 = op1(e13, e10)) | ~ spl144_15), inference(avatar_component_clause, [], [f1230])).
fof(f1230, plain, (spl144_15 <=> (e12 = op1(e13, e10))), introduced(avatar_definition, [new_symbols(naming, [spl144_15])])).
fof(f3226, plain, (~ (op2(e23, e20) = h4(op1(e13, e10))) | spl144_325), inference(avatar_component_clause, [], [f3224])).
fof(f3224, plain, (spl144_325 <=> (op2(e23, e20) = h4(op1(e13, e10)))), introduced(avatar_definition, [new_symbols(naming, [spl144_325])])).
fof(f4025, plain, (~ spl144_47 | ~ spl144_35), inference(avatar_split_clause, [], [f4024, f1315, f1366])).
fof(f1366, plain, (spl144_47 <=> (e12 = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl144_47])])).
fof(f4024, plain, (~ (e12 = op1(e11, e10)) | ~ spl144_35), inference(forward_demodulation, [], [f440, f1317])).
fof(f440, plain, ~ (op1(e11, e10) = op1(e11, e13)), inference(cnf_transformation, [], [f5])).
fof(f4023, plain, (~ spl144_37 | ~ spl144_5), inference(avatar_split_clause, [], [f4022, f1188, f1324])).
fof(f1188, plain, (spl144_5 <=> (e10 = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl144_5])])).
fof(f4022, plain, (~ (e10 = op1(e11, e12)) | ~ spl144_5), inference(forward_demodulation, [], [f424, f1190])).
fof(f1190, plain, ((e10 = op1(e13, e12)) | ~ spl144_5), inference(avatar_component_clause, [], [f1188])).
fof(f424, plain, ~ (op1(e11, e12) = op1(e13, e12)), inference(cnf_transformation, [], [f5])).
fof(f3988, plain, (~ spl144_10 | ~ spl144_2), inference(avatar_split_clause, [], [f3987, f1175, f1209])).
fof(f1209, plain, (spl144_10 <=> (e11 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl144_10])])).
fof(f3987, plain, (~ (e11 = op1(e13, e11)) | ~ spl144_2), inference(forward_demodulation, [], [f454, f1177])).
fof(f454, plain, ~ (op1(e13, e11) = op1(e13, e13)), inference(cnf_transformation, [], [f5])).
fof(f3981, plain, (spl144_35 | ~ spl144_2), inference(avatar_split_clause, [], [f3980, f1175, f1315])).
fof(f3980, plain, ((e12 = op1(e11, e13)) | ~ spl144_2), inference(forward_demodulation, [], [f1086, f1177])).
fof(f1086, plain, (e12 = op1(op1(e13, e13), e13)), inference(cnf_transformation, [], [f12])).
fof(f12, plain, ((e12 = op1(op1(e13, e13), e13)) & (e11 = op1(e13, e13)) & (e10 = op1(op1(e13, e13), op1(op1(e13, e13), e13)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG115+1.p', ax12)).
fof(f3977, plain, (spl144_37 | ~ spl144_2), inference(avatar_split_clause, [], [f3976, f1175, f1324])).
fof(f3976, plain, ((e10 = op1(e11, e12)) | ~ spl144_2), inference(forward_demodulation, [], [f2987, f1177])).
fof(f2987, plain, (e10 = op1(op1(e13, e13), e12)), inference(forward_demodulation, [], [f1084, f1086])).
fof(f1084, plain, (e10 = op1(op1(e13, e13), op1(op1(e13, e13), e13))), inference(cnf_transformation, [], [f12])).
fof(f3975, plain, (~ spl144_2 | ~ spl144_3), inference(avatar_contradiction_clause, [], [f3974])).
fof(f3974, plain, ($false | (~ spl144_2 | ~ spl144_3)), inference(subsumption_resolution, [], [f3973, f507])).
fof(f3973, plain, ((e11 = e12) | (~ spl144_2 | ~ spl144_3)), inference(backward_demodulation, [], [f1181, f1177])).
fof(f1181, plain, ((e12 = op1(e13, e13)) | ~ spl144_3), inference(avatar_component_clause, [], [f1179])).
fof(f1179, plain, (spl144_3 <=> (e12 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl144_3])])).
fof(f3967, plain, (spl144_3 | ~ spl144_4), inference(avatar_split_clause, [], [f3964, f1183, f1179])).
fof(f1183, plain, (spl144_4 <=> (e13 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl144_4])])).
fof(f3964, plain, ((e12 = op1(e13, e13)) | ~ spl144_4), inference(backward_demodulation, [], [f1086, f1185])).
fof(f1185, plain, ((e13 = op1(e13, e13)) | ~ spl144_4), inference(avatar_component_clause, [], [f1183])).
fof(f3935, plain, (~ spl144_15 | ~ spl144_16), inference(avatar_contradiction_clause, [], [f3934])).
fof(f3934, plain, ($false | (~ spl144_15 | ~ spl144_16)), inference(subsumption_resolution, [], [f3933, f509])).
fof(f3933, plain, ((e12 = e13) | (~ spl144_15 | ~ spl144_16)), inference(backward_demodulation, [], [f1236, f1232])).
fof(f1236, plain, ((e13 = op1(e13, e10)) | ~ spl144_16), inference(avatar_component_clause, [], [f1234])).
fof(f1234, plain, (spl144_16 <=> (e13 = op1(e13, e10))), introduced(avatar_definition, [new_symbols(naming, [spl144_16])])).
fof(f3904, plain, (~ spl144_9 | ~ spl144_25), inference(avatar_split_clause, [], [f3901, f1273, f1205])).
fof(f1273, plain, (spl144_25 <=> (e10 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl144_25])])).
fof(f3901, plain, (~ (e10 = op1(e13, e11)) | ~ spl144_25), inference(backward_demodulation, [], [f419, f1275])).
fof(f1275, plain, ((e10 = op1(e12, e11)) | ~ spl144_25), inference(avatar_component_clause, [], [f1273])).
fof(f419, plain, ~ (op1(e12, e11) = op1(e13, e11)), inference(cnf_transformation, [], [f5])).
fof(f3889, plain, (~ spl144_35 | ~ spl144_36), inference(avatar_contradiction_clause, [], [f3888])).
fof(f3888, plain, ($false | (~ spl144_35 | ~ spl144_36)), inference(subsumption_resolution, [], [f3887, f509])).
fof(f3887, plain, ((e12 = e13) | (~ spl144_35 | ~ spl144_36)), inference(backward_demodulation, [], [f1321, f1317])).
fof(f1321, plain, ((e13 = op1(e11, e13)) | ~ spl144_36), inference(avatar_component_clause, [], [f1319])).
fof(f1319, plain, (spl144_36 <=> (e13 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl144_36])])).
fof(f3878, plain, (~ spl144_21 | ~ spl144_37), inference(avatar_split_clause, [], [f3874, f1324, f1256])).
fof(f1256, plain, (spl144_21 <=> (e10 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl144_21])])).
fof(f3874, plain, (~ (e10 = op1(e12, e12)) | ~ spl144_37), inference(backward_demodulation, [], [f423, f1326])).
fof(f423, plain, ~ (op1(e11, e12) = op1(e12, e12)), inference(cnf_transformation, [], [f5])).
fof(f3872, plain, (~ spl144_37 | ~ spl144_41), inference(avatar_split_clause, [], [f3868, f1341, f1324])).
fof(f1341, plain, (spl144_41 <=> (e10 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl144_41])])).
fof(f3868, plain, (~ (e10 = op1(e11, e12)) | ~ spl144_41), inference(backward_demodulation, [], [f441, f1343])).
fof(f1343, plain, ((e10 = op1(e11, e11)) | ~ spl144_41), inference(avatar_component_clause, [], [f1341])).
fof(f441, plain, ~ (op1(e11, e11) = op1(e11, e12)), inference(cnf_transformation, [], [f5])).
fof(f3864, plain, (~ spl144_37 | ~ spl144_45), inference(avatar_split_clause, [], [f3859, f1358, f1324])).
fof(f1358, plain, (spl144_45 <=> (e10 = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl144_45])])).
fof(f3859, plain, (~ (e10 = op1(e11, e12)) | ~ spl144_45), inference(backward_demodulation, [], [f439, f1360])).
fof(f1360, plain, ((e10 = op1(e11, e10)) | ~ spl144_45), inference(avatar_component_clause, [], [f1358])).
fof(f439, plain, ~ (op1(e11, e10) = op1(e11, e12)), inference(cnf_transformation, [], [f5])).
fof(f3822, plain, (~ spl144_13 | ~ spl144_61), inference(avatar_split_clause, [], [f3814, f1426, f1222])).
fof(f1222, plain, (spl144_13 <=> (e10 = op1(e13, e10))), introduced(avatar_definition, [new_symbols(naming, [spl144_13])])).
fof(f3814, plain, (~ (e10 = op1(e13, e10)) | ~ spl144_61), inference(backward_demodulation, [], [f410, f1428])).
fof(f410, plain, ~ (op1(e10, e10) = op1(e13, e10)), inference(cnf_transformation, [], [f5])).
fof(f3821, plain, (~ spl144_29 | ~ spl144_61), inference(avatar_split_clause, [], [f3813, f1426, f1290])).
fof(f1290, plain, (spl144_29 <=> (e10 = op1(e12, e10))), introduced(avatar_definition, [new_symbols(naming, [spl144_29])])).
fof(f3813, plain, (~ (e10 = op1(e12, e10)) | ~ spl144_61), inference(backward_demodulation, [], [f409, f1428])).
fof(f409, plain, ~ (op1(e10, e10) = op1(e12, e10)), inference(cnf_transformation, [], [f5])).
fof(f3811, plain, (~ spl144_65 | ~ spl144_66), inference(avatar_contradiction_clause, [], [f3810])).
fof(f3810, plain, ($false | (~ spl144_65 | ~ spl144_66)), inference(subsumption_resolution, [], [f3809, f510])).
fof(f3809, plain, ((e20 = e21) | (~ spl144_65 | ~ spl144_66)), inference(backward_demodulation, [], [f1481, f1477])).
fof(f1477, plain, ((e20 = op2(e23, e23)) | ~ spl144_65), inference(avatar_component_clause, [], [f1475])).
fof(f1475, plain, (spl144_65 <=> (e20 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl144_65])])).
fof(f3808, plain, (~ spl144_66 | ~ spl144_67), inference(avatar_contradiction_clause, [], [f3807])).
fof(f3807, plain, ($false | (~ spl144_66 | ~ spl144_67)), inference(subsumption_resolution, [], [f3806, f513])).
fof(f3806, plain, ((e21 = e22) | (~ spl144_66 | ~ spl144_67)), inference(backward_demodulation, [], [f1485, f1481])).
fof(f1485, plain, ((e22 = op2(e23, e23)) | ~ spl144_67), inference(avatar_component_clause, [], [f1483])).
fof(f1483, plain, (spl144_67 <=> (e22 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl144_67])])).
fof(f3749, plain, (spl144_303 | ~ spl144_85), inference(avatar_split_clause, [], [f3748, f1560, f3088])).
fof(f3088, plain, (spl144_303 <=> (e20 = h3(e11))), introduced(avatar_definition, [new_symbols(naming, [spl144_303])])).
fof(f1560, plain, (spl144_85 <=> (e20 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl144_85])])).
fof(f3748, plain, ((e20 = h3(e11)) | ~ spl144_85), inference(backward_demodulation, [], [f1100, f1562])).
fof(f1562, plain, ((e20 = op2(e22, e22)) | ~ spl144_85), inference(avatar_component_clause, [], [f1560])).
fof(f3746, plain, (~ spl144_73 | ~ spl144_89), inference(avatar_split_clause, [], [f3742, f1577, f1509])).
fof(f1577, plain, (spl144_89 <=> (e20 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl144_89])])).
fof(f3742, plain, (~ (e20 = op2(e23, e21)) | ~ spl144_89), inference(backward_demodulation, [], [f467, f1579])).
fof(f1579, plain, ((e20 = op2(e22, e21)) | ~ spl144_89), inference(avatar_component_clause, [], [f1577])).
fof(f467, plain, ~ (op2(e22, e21) = op2(e23, e21)), inference(cnf_transformation, [], [f6])).
fof(f3718, plain, (~ spl144_69 | ~ spl144_101), inference(avatar_split_clause, [], [f3714, f1628, f1492])).
fof(f1492, plain, (spl144_69 <=> (e20 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl144_69])])).
fof(f3714, plain, (~ (e20 = op2(e23, e22)) | ~ spl144_101), inference(backward_demodulation, [], [f472, f1630])).
fof(f472, plain, ~ (op2(e21, e22) = op2(e23, e22)), inference(cnf_transformation, [], [f6])).
fof(f3717, plain, (~ spl144_311 | ~ spl144_101), inference(avatar_split_clause, [], [f3713, f1628, f3129])).
fof(f3129, plain, (spl144_311 <=> (e20 = h2(e11))), introduced(avatar_definition, [new_symbols(naming, [spl144_311])])).
fof(f3713, plain, (~ (e20 = h2(e11)) | ~ spl144_101), inference(backward_demodulation, [], [f3003, f1630])).
fof(f3003, plain, ~ (op2(e21, e22) = h2(e11)), inference(backward_demodulation, [], [f489, f1096])).
fof(f489, plain, ~ (op2(e21, e21) = op2(e21, e22)), inference(cnf_transformation, [], [f6])).
fof(f3716, plain, (~ spl144_303 | ~ spl144_101), inference(avatar_split_clause, [], [f3712, f1628, f3088])).
fof(f3712, plain, (~ (e20 = h3(e11)) | ~ spl144_101), inference(backward_demodulation, [], [f3009, f1630])).
fof(f3009, plain, ~ (op2(e21, e22) = h3(e11)), inference(backward_demodulation, [], [f471, f1100])).
fof(f471, plain, ~ (op2(e21, e22) = op2(e22, e22)), inference(cnf_transformation, [], [f6])).
fof(f3711, plain, (spl144_311 | ~ spl144_105), inference(avatar_split_clause, [], [f3710, f1645, f3129])).
fof(f1645, plain, (spl144_105 <=> (e20 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl144_105])])).
fof(f3710, plain, ((e20 = h2(e11)) | ~ spl144_105), inference(backward_demodulation, [], [f1096, f1647])).
fof(f1647, plain, ((e20 = op2(e21, e21)) | ~ spl144_105), inference(avatar_component_clause, [], [f1645])).
fof(f3693, plain, (~ spl144_84 | ~ spl144_116), inference(avatar_split_clause, [], [f3690, f1691, f1555])).
fof(f1555, plain, (spl144_84 <=> (e23 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl144_84])])).
fof(f3690, plain, (~ (e23 = op2(e22, e23)) | ~ spl144_116), inference(backward_demodulation, [], [f475, f1693])).
fof(f475, plain, ~ (op2(e20, e23) = op2(e22, e23)), inference(cnf_transformation, [], [f6])).
fof(f3662, plain, (~ spl144_77 | ~ spl144_125), inference(avatar_split_clause, [], [f3654, f1730, f1526])).
fof(f1526, plain, (spl144_77 <=> (e20 = op2(e23, e20))), introduced(avatar_definition, [new_symbols(naming, [spl144_77])])).
fof(f3654, plain, (~ (e20 = op2(e23, e20)) | ~ spl144_125), inference(backward_demodulation, [], [f2992, f3651])).
fof(f3651, plain, ((e20 = h1(e11)) | ~ spl144_125), inference(backward_demodulation, [], [f1092, f1732])).
fof(f2992, plain, ~ (op2(e23, e20) = h1(e11)), inference(backward_demodulation, [], [f458, f1092])).
fof(f458, plain, ~ (op2(e20, e20) = op2(e23, e20)), inference(cnf_transformation, [], [f6])).
fof(f3661, plain, (~ spl144_93 | ~ spl144_125), inference(avatar_split_clause, [], [f3653, f1730, f1594])).
fof(f1594, plain, (spl144_93 <=> (e20 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl144_93])])).
fof(f3653, plain, (~ (e20 = op2(e22, e20)) | ~ spl144_125), inference(backward_demodulation, [], [f2991, f3651])).
fof(f2991, plain, ~ (op2(e22, e20) = h1(e11)), inference(backward_demodulation, [], [f457, f1092])).
fof(f457, plain, ~ (op2(e20, e20) = op2(e22, e20)), inference(cnf_transformation, [], [f6])).
fof(f3271, plain, (~ spl144_321 | spl144_295 | spl144_293 | ~ spl144_322 | ~ spl144_323 | ~ spl144_324 | ~ spl144_325 | ~ spl144_326 | ~ spl144_327 | ~ spl144_328 | ~ spl144_329 | ~ spl144_330 | ~ spl144_331 | ~ spl144_332 | ~ spl144_333 | ~ spl144_334 | ~ spl144_335 | ~ spl144_336), inference(avatar_split_clause, [], [f3206, f3268, f3264, f3260, f3256, f3252, f3248, f3244, f3240, f3236, f3232, f3228, f3224, f3220, f3216, f3212, f3034, f3046, f3208])).
fof(f3046, plain, (spl144_295 <=> sP141), introduced(avatar_definition, [new_symbols(naming, [spl144_295])])).
fof(f3034, plain, (spl144_293 <=> sP142), introduced(avatar_definition, [new_symbols(naming, [spl144_293])])).
fof(f3206, plain, (~ (h1(e11) = h4(op1(e10, e10))) | ~ (h4(op1(e10, e11)) = op2(e20, h4(e11))) | ~ (op2(e20, e22) = h4(op1(e10, e12))) | ~ (op2(e20, e23) = h4(op1(e10, e13))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), e20)) | ~ (e20 = h4(op1(e11, e12))) | ~ (e22 = h4(op1(e11, e13))) | ~ (op2(e22, e20) = h4(op1(e12, e10))) | ~ (h4(op1(e12, e11)) = op2(e22, h4(e11))) | ~ (h3(e11) = h4(op1(e12, e12))) | ~ (op2(e22, e23) = h4(op1(e12, e13))) | ~ (op2(e23, e20) = h4(op1(e13, e10))) | ~ (h4(op1(e13, e11)) = op2(e23, h4(e11))) | ~ (op2(e23, e22) = h4(op1(e13, e12))) | ~ (h4(e11) = h4(op1(e13, e13))) | sP142 | sP141 | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11)))), inference(forward_demodulation, [], [f3205, f1092])).
fof(f3205, plain, (~ (op2(e20, e20) = h4(op1(e10, e10))) | ~ (h4(op1(e10, e11)) = op2(e20, h4(e11))) | ~ (op2(e20, e22) = h4(op1(e10, e12))) | ~ (op2(e20, e23) = h4(op1(e10, e13))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), e20)) | ~ (e20 = h4(op1(e11, e12))) | ~ (e22 = h4(op1(e11, e13))) | ~ (op2(e22, e20) = h4(op1(e12, e10))) | ~ (h4(op1(e12, e11)) = op2(e22, h4(e11))) | ~ (h3(e11) = h4(op1(e12, e12))) | ~ (op2(e22, e23) = h4(op1(e12, e13))) | ~ (op2(e23, e20) = h4(op1(e13, e10))) | ~ (h4(op1(e13, e11)) = op2(e23, h4(e11))) | ~ (op2(e23, e22) = h4(op1(e13, e12))) | ~ (h4(e11) = h4(op1(e13, e13))) | sP142 | sP141 | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11)))), inference(forward_demodulation, [], [f3204, f3028])).
fof(f3204, plain, (~ (h4(op1(e10, e11)) = op2(e20, h4(e11))) | ~ (op2(e20, e22) = h4(op1(e10, e12))) | ~ (op2(e20, e23) = h4(op1(e10, e13))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), e20)) | ~ (e20 = h4(op1(e11, e12))) | ~ (e22 = h4(op1(e11, e13))) | ~ (op2(e22, e20) = h4(op1(e12, e10))) | ~ (h4(op1(e12, e11)) = op2(e22, h4(e11))) | ~ (h3(e11) = h4(op1(e12, e12))) | ~ (op2(e22, e23) = h4(op1(e12, e13))) | ~ (op2(e23, e20) = h4(op1(e13, e10))) | ~ (h4(op1(e13, e11)) = op2(e23, h4(e11))) | ~ (op2(e23, e22) = h4(op1(e13, e12))) | ~ (h4(e11) = h4(op1(e13, e13))) | sP142 | sP141 | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f3203, f3028])).
fof(f3203, plain, (~ (op2(e20, e22) = h4(op1(e10, e12))) | ~ (op2(e20, e23) = h4(op1(e10, e13))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), e20)) | ~ (e20 = h4(op1(e11, e12))) | ~ (e22 = h4(op1(e11, e13))) | ~ (op2(e22, e20) = h4(op1(e12, e10))) | ~ (h4(op1(e12, e11)) = op2(e22, h4(e11))) | ~ (h3(e11) = h4(op1(e12, e12))) | ~ (op2(e22, e23) = h4(op1(e12, e13))) | ~ (op2(e23, e20) = h4(op1(e13, e10))) | ~ (h4(op1(e13, e11)) = op2(e23, h4(e11))) | ~ (op2(e23, e22) = h4(op1(e13, e12))) | ~ (h4(e11) = h4(op1(e13, e13))) | sP142 | sP141 | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f3202, f3028])).
fof(f3202, plain, (~ (h4(op1(e10, e12)) = op2(h4(e10), e22)) | ~ (op2(e20, e23) = h4(op1(e10, e13))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), e20)) | ~ (e20 = h4(op1(e11, e12))) | ~ (e22 = h4(op1(e11, e13))) | ~ (op2(e22, e20) = h4(op1(e12, e10))) | ~ (h4(op1(e12, e11)) = op2(e22, h4(e11))) | ~ (h3(e11) = h4(op1(e12, e12))) | ~ (op2(e22, e23) = h4(op1(e12, e13))) | ~ (op2(e23, e20) = h4(op1(e13, e10))) | ~ (h4(op1(e13, e11)) = op2(e23, h4(e11))) | ~ (op2(e23, e22) = h4(op1(e13, e12))) | ~ (h4(e11) = h4(op1(e13, e13))) | sP142 | sP141 | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f3201, f3017])).
fof(f3201, plain, (~ (op2(e20, e23) = h4(op1(e10, e13))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), e20)) | ~ (e20 = h4(op1(e11, e12))) | ~ (e22 = h4(op1(e11, e13))) | ~ (op2(e22, e20) = h4(op1(e12, e10))) | ~ (h4(op1(e12, e11)) = op2(e22, h4(e11))) | ~ (h3(e11) = h4(op1(e12, e12))) | ~ (op2(e22, e23) = h4(op1(e12, e13))) | ~ (op2(e23, e20) = h4(op1(e13, e10))) | ~ (h4(op1(e13, e11)) = op2(e23, h4(e11))) | ~ (op2(e23, e22) = h4(op1(e13, e12))) | ~ (h4(e11) = h4(op1(e13, e13))) | sP142 | sP141 | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f3200, f3028])).
fof(f3200, plain, (~ (h4(op1(e10, e13)) = op2(h4(e10), e23)) | ~ (h4(op1(e11, e10)) = op2(h4(e11), e20)) | ~ (e20 = h4(op1(e11, e12))) | ~ (e22 = h4(op1(e11, e13))) | ~ (op2(e22, e20) = h4(op1(e12, e10))) | ~ (h4(op1(e12, e11)) = op2(e22, h4(e11))) | ~ (h3(e11) = h4(op1(e12, e12))) | ~ (op2(e22, e23) = h4(op1(e12, e13))) | ~ (op2(e23, e20) = h4(op1(e13, e10))) | ~ (h4(op1(e13, e11)) = op2(e23, h4(e11))) | ~ (op2(e23, e22) = h4(op1(e13, e12))) | ~ (h4(e11) = h4(op1(e13, e13))) | sP142 | sP141 | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f3199, f1102])).
fof(f3199, plain, (~ (h4(op1(e11, e10)) = op2(h4(e11), e20)) | ~ (e20 = h4(op1(e11, e12))) | ~ (e22 = h4(op1(e11, e13))) | ~ (op2(e22, e20) = h4(op1(e12, e10))) | ~ (h4(op1(e12, e11)) = op2(e22, h4(e11))) | ~ (h3(e11) = h4(op1(e12, e12))) | ~ (op2(e22, e23) = h4(op1(e12, e13))) | ~ (op2(e23, e20) = h4(op1(e13, e10))) | ~ (h4(op1(e13, e11)) = op2(e23, h4(e11))) | ~ (op2(e23, e22) = h4(op1(e13, e12))) | ~ (h4(e11) = h4(op1(e13, e13))) | sP142 | sP141 | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f3198, f3028])).
fof(f3198, plain, (~ (e20 = h4(op1(e11, e12))) | ~ (e22 = h4(op1(e11, e13))) | ~ (op2(e22, e20) = h4(op1(e12, e10))) | ~ (h4(op1(e12, e11)) = op2(e22, h4(e11))) | ~ (h3(e11) = h4(op1(e12, e12))) | ~ (op2(e22, e23) = h4(op1(e12, e13))) | ~ (op2(e23, e20) = h4(op1(e13, e10))) | ~ (h4(op1(e13, e11)) = op2(e23, h4(e11))) | ~ (op2(e23, e22) = h4(op1(e13, e12))) | ~ (h4(e11) = h4(op1(e13, e13))) | sP142 | sP141 | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f3197, f3025])).
fof(f3197, plain, (~ (h4(op1(e11, e12)) = op2(h4(e11), e22)) | ~ (e22 = h4(op1(e11, e13))) | ~ (op2(e22, e20) = h4(op1(e12, e10))) | ~ (h4(op1(e12, e11)) = op2(e22, h4(e11))) | ~ (h3(e11) = h4(op1(e12, e12))) | ~ (op2(e22, e23) = h4(op1(e12, e13))) | ~ (op2(e23, e20) = h4(op1(e13, e10))) | ~ (h4(op1(e13, e11)) = op2(e23, h4(e11))) | ~ (op2(e23, e22) = h4(op1(e13, e12))) | ~ (h4(e11) = h4(op1(e13, e13))) | sP142 | sP141 | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f3196, f3017])).
fof(f3196, plain, (~ (e22 = h4(op1(e11, e13))) | ~ (op2(e22, e20) = h4(op1(e12, e10))) | ~ (h4(op1(e12, e11)) = op2(e22, h4(e11))) | ~ (h3(e11) = h4(op1(e12, e12))) | ~ (op2(e22, e23) = h4(op1(e12, e13))) | ~ (op2(e23, e20) = h4(op1(e13, e10))) | ~ (h4(op1(e13, e11)) = op2(e23, h4(e11))) | ~ (op2(e23, e22) = h4(op1(e13, e12))) | ~ (h4(e11) = h4(op1(e13, e13))) | sP142 | sP141 | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f3195, f3024])).
fof(f3195, plain, (~ (h4(op1(e11, e13)) = op2(h4(e11), e23)) | ~ (op2(e22, e20) = h4(op1(e12, e10))) | ~ (h4(op1(e12, e11)) = op2(e22, h4(e11))) | ~ (h3(e11) = h4(op1(e12, e12))) | ~ (op2(e22, e23) = h4(op1(e12, e13))) | ~ (op2(e23, e20) = h4(op1(e13, e10))) | ~ (h4(op1(e13, e11)) = op2(e23, h4(e11))) | ~ (op2(e23, e22) = h4(op1(e13, e12))) | ~ (h4(e11) = h4(op1(e13, e13))) | sP142 | sP141 | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f3194, f1102])).
fof(f3194, plain, (~ (op2(e22, e20) = h4(op1(e12, e10))) | ~ (h4(op1(e12, e11)) = op2(e22, h4(e11))) | ~ (h3(e11) = h4(op1(e12, e12))) | ~ (op2(e22, e23) = h4(op1(e12, e13))) | ~ (op2(e23, e20) = h4(op1(e13, e10))) | ~ (h4(op1(e13, e11)) = op2(e23, h4(e11))) | ~ (op2(e23, e22) = h4(op1(e13, e12))) | ~ (h4(e11) = h4(op1(e13, e13))) | sP142 | sP141 | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f3193, f3017])).
fof(f3193, plain, (~ (h4(op1(e12, e10)) = op2(h4(e12), e20)) | ~ (h4(op1(e12, e11)) = op2(e22, h4(e11))) | ~ (h3(e11) = h4(op1(e12, e12))) | ~ (op2(e22, e23) = h4(op1(e12, e13))) | ~ (op2(e23, e20) = h4(op1(e13, e10))) | ~ (h4(op1(e13, e11)) = op2(e23, h4(e11))) | ~ (op2(e23, e22) = h4(op1(e13, e12))) | ~ (h4(e11) = h4(op1(e13, e13))) | sP142 | sP141 | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f3192, f3028])).
fof(f3192, plain, (~ (h4(op1(e12, e11)) = op2(e22, h4(e11))) | ~ (h3(e11) = h4(op1(e12, e12))) | ~ (op2(e22, e23) = h4(op1(e12, e13))) | ~ (op2(e23, e20) = h4(op1(e13, e10))) | ~ (h4(op1(e13, e11)) = op2(e23, h4(e11))) | ~ (op2(e23, e22) = h4(op1(e13, e12))) | ~ (h4(e11) = h4(op1(e13, e13))) | sP142 | sP141 | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f3191, f3017])).
fof(f3191, plain, (~ (h3(e11) = h4(op1(e12, e12))) | ~ (op2(e22, e23) = h4(op1(e12, e13))) | ~ (op2(e23, e20) = h4(op1(e13, e10))) | ~ (h4(op1(e13, e11)) = op2(e23, h4(e11))) | ~ (op2(e23, e22) = h4(op1(e13, e12))) | ~ (h4(e11) = h4(op1(e13, e13))) | sP142 | sP141 | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f3190, f1100])).
fof(f3190, plain, (~ (op2(e22, e22) = h4(op1(e12, e12))) | ~ (op2(e22, e23) = h4(op1(e12, e13))) | ~ (op2(e23, e20) = h4(op1(e13, e10))) | ~ (h4(op1(e13, e11)) = op2(e23, h4(e11))) | ~ (op2(e23, e22) = h4(op1(e13, e12))) | ~ (h4(e11) = h4(op1(e13, e13))) | sP142 | sP141 | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f3189, f3017])).
fof(f3189, plain, (~ (op2(e22, e23) = h4(op1(e12, e13))) | ~ (op2(e23, e20) = h4(op1(e13, e10))) | ~ (h4(op1(e13, e11)) = op2(e23, h4(e11))) | ~ (op2(e23, e22) = h4(op1(e13, e12))) | ~ (h4(e11) = h4(op1(e13, e13))) | sP142 | sP141 | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f3188, f3017])).
fof(f3188, plain, (~ (h4(op1(e12, e13)) = op2(h4(e12), e23)) | ~ (op2(e23, e20) = h4(op1(e13, e10))) | ~ (h4(op1(e13, e11)) = op2(e23, h4(e11))) | ~ (op2(e23, e22) = h4(op1(e13, e12))) | ~ (h4(e11) = h4(op1(e13, e13))) | sP142 | sP141 | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f3187, f1102])).
fof(f3187, plain, (~ (op2(e23, e20) = h4(op1(e13, e10))) | ~ (h4(op1(e13, e11)) = op2(e23, h4(e11))) | ~ (op2(e23, e22) = h4(op1(e13, e12))) | ~ (h4(e11) = h4(op1(e13, e13))) | sP142 | sP141 | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f3186, f1102])).
fof(f3186, plain, (~ (h4(op1(e13, e10)) = op2(h4(e13), e20)) | ~ (h4(op1(e13, e11)) = op2(e23, h4(e11))) | ~ (op2(e23, e22) = h4(op1(e13, e12))) | ~ (h4(e11) = h4(op1(e13, e13))) | sP142 | sP141 | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f3185, f3028])).
fof(f3185, plain, (~ (h4(op1(e13, e11)) = op2(e23, h4(e11))) | ~ (op2(e23, e22) = h4(op1(e13, e12))) | ~ (h4(e11) = h4(op1(e13, e13))) | sP142 | sP141 | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f3184, f1102])).
fof(f3184, plain, (~ (op2(e23, e22) = h4(op1(e13, e12))) | ~ (h4(e11) = h4(op1(e13, e13))) | sP142 | sP141 | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f3183, f1102])).
fof(f3183, plain, (~ (h4(op1(e13, e12)) = op2(h4(e13), e22)) | ~ (h4(e11) = h4(op1(e13, e13))) | sP142 | sP141 | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f3182, f3017])).
fof(f3182, plain, (~ (h4(e11) = h4(op1(e13, e13))) | sP142 | sP141 | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f3181, f1104])).
fof(f3181, plain, (~ (op2(e23, e23) = h4(op1(e13, e13))) | sP142 | sP141 | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f3180, f1102])).
fof(f3180, plain, (sP142 | sP141 | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(subsumption_resolution, [], [f3179, f3030])).
fof(f3030, plain, ~ sP143, inference(subsumption_resolution, [], [f1108, f3017])).
fof(f1108, plain, (~ (e22 = h4(e12)) | ~ sP143), inference(cnf_transformation, [], [f300])).
fof(f300, plain, ((~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10))) | ~ sP143), inference(nnf_transformation, [], [f166])).
fof(f166, plain, ((~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10))) | ~ sP143), inference(usedef, [], [e166])).
fof(e166, plain, (sP143 <=> (~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP143])])).
fof(f3179, plain, (sP143 | sP142 | sP141 | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(subsumption_resolution, [], [f1169, f1102])).
fof(f1169, plain, (~ (e23 = h4(e13)) | sP143 | sP142 | sP141 | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(cnf_transformation, [], [f167])).
fof(f167, plain, (((~ (e23 = h4(e13)) & ~ (e23 = h4(e12)) & ~ (e23 = h4(e11)) & ~ (e23 = h4(e10))) | sP143 | sP142 | sP141 | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) & ((~ (e23 = h3(e13)) & ~ (e23 = h3(e12)) & ~ (e23 = h3(e11)) & ~ (e23 = h3(e10))) | sP140 | sP139 | sP138 | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) & ((~ (e23 = h2(e13)) & ~ (e23 = h2(e12)) & ~ (e23 = h2(e11)) & ~ (e23 = h2(e10))) | sP137 | sP136 | sP135 | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) & ((~ (e23 = h1(e13)) & ~ (e23 = h1(e12)) & ~ (e23 = h1(e11)) & ~ (e23 = h1(e10))) | sP134 | sP133 | sP132 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(definition_folding, [], [f20, e166, e165, e164, e163, e162, e161, e160, e159, e158, e157, e156, e155])).
fof(f155, plain, ((~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10))) | ~ sP132), inference(usedef, [], [e155])).
fof(e155, plain, (sP132 <=> (~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP132])])).
fof(f156, plain, ((~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10))) | ~ sP133), inference(usedef, [], [e156])).
fof(e156, plain, (sP133 <=> (~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP133])])).
fof(f157, plain, ((~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10))) | ~ sP134), inference(usedef, [], [e157])).
fof(e157, plain, (sP134 <=> (~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP134])])).
fof(f158, plain, ((~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10))) | ~ sP135), inference(usedef, [], [e158])).
fof(e158, plain, (sP135 <=> (~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP135])])).
fof(f159, plain, ((~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10))) | ~ sP136), inference(usedef, [], [e159])).
fof(e159, plain, (sP136 <=> (~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP136])])).
fof(f160, plain, ((~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10))) | ~ sP137), inference(usedef, [], [e160])).
fof(e160, plain, (sP137 <=> (~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP137])])).
fof(f161, plain, ((~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10))) | ~ sP138), inference(usedef, [], [e161])).
fof(e161, plain, (sP138 <=> (~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP138])])).
fof(f162, plain, ((~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10))) | ~ sP139), inference(usedef, [], [e162])).
fof(e162, plain, (sP139 <=> (~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP139])])).
fof(f163, plain, ((~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10))) | ~ sP140), inference(usedef, [], [e163])).
fof(e163, plain, (sP140 <=> (~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP140])])).
fof(f164, plain, ((~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10))) | ~ sP141), inference(usedef, [], [e164])).
fof(e164, plain, (sP141 <=> (~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP141])])).
fof(f165, plain, ((~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10))) | ~ sP142), inference(usedef, [], [e165])).
fof(e165, plain, (sP142 <=> (~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP142])])).
fof(f20, plain, (((~ (e23 = h4(e13)) & ~ (e23 = h4(e12)) & ~ (e23 = h4(e11)) & ~ (e23 = h4(e10))) | (~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10))) | (~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10))) | (~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10))) | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) & ((~ (e23 = h3(e13)) & ~ (e23 = h3(e12)) & ~ (e23 = h3(e11)) & ~ (e23 = h3(e10))) | (~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10))) | (~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10))) | (~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10))) | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) & ((~ (e23 = h2(e13)) & ~ (e23 = h2(e12)) & ~ (e23 = h2(e11)) & ~ (e23 = h2(e10))) | (~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10))) | (~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10))) | (~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) & ((~ (e23 = h1(e13)) & ~ (e23 = h1(e12)) & ~ (e23 = h1(e11)) & ~ (e23 = h1(e10))) | (~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10))) | (~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10))) | (~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10))) | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(ennf_transformation, [], [f19])).
fof(f19, plain, ~ ((((e23 = h4(e13)) | (e23 = h4(e12)) | (e23 = h4(e11)) | (e23 = h4(e10))) & ((e22 = h4(e13)) | (e22 = h4(e12)) | (e22 = h4(e11)) | (e22 = h4(e10))) & ((e21 = h4(e13)) | (e21 = h4(e12)) | (e21 = h4(e11)) | (e21 = h4(e10))) & ((e20 = h4(e13)) | (e20 = h4(e12)) | (e20 = h4(e11)) | (e20 = h4(e10))) & (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) & (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) & (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) & (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) & (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) & (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) & (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) & (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) & (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) & (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) & (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) & (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) & (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) & (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) & (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) & (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) | (((e23 = h3(e13)) | (e23 = h3(e12)) | (e23 = h3(e11)) | (e23 = h3(e10))) & ((e22 = h3(e13)) | (e22 = h3(e12)) | (e22 = h3(e11)) | (e22 = h3(e10))) & ((e21 = h3(e13)) | (e21 = h3(e12)) | (e21 = h3(e11)) | (e21 = h3(e10))) & ((e20 = h3(e13)) | (e20 = h3(e12)) | (e20 = h3(e11)) | (e20 = h3(e10))) & (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) & (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) & (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) & (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) & (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) & (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) & (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) & (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) & (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) & (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) & (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) & (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) & (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) & (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) & (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) & (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) | (((e23 = h2(e13)) | (e23 = h2(e12)) | (e23 = h2(e11)) | (e23 = h2(e10))) & ((e22 = h2(e13)) | (e22 = h2(e12)) | (e22 = h2(e11)) | (e22 = h2(e10))) & ((e21 = h2(e13)) | (e21 = h2(e12)) | (e21 = h2(e11)) | (e21 = h2(e10))) & ((e20 = h2(e13)) | (e20 = h2(e12)) | (e20 = h2(e11)) | (e20 = h2(e10))) & (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) & (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) & (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) & (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) & (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) & (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) & (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) & (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) & (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) & (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) & (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) & (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) & (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) & (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) & (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) & (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) | (((e23 = h1(e13)) | (e23 = h1(e12)) | (e23 = h1(e11)) | (e23 = h1(e10))) & ((e22 = h1(e13)) | (e22 = h1(e12)) | (e22 = h1(e11)) | (e22 = h1(e10))) & ((e21 = h1(e13)) | (e21 = h1(e12)) | (e21 = h1(e11)) | (e21 = h1(e10))) & ((e20 = h1(e13)) | (e20 = h1(e12)) | (e20 = h1(e11)) | (e20 = h1(e10))) & (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) & (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) & (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) & (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) & (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) & (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) & (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) & (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) & (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) & (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) & (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) & (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) & (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) & (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) & (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) & (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(negated_conjecture, [], [f18])).
fof(f18, plain, ~ ((((e23 = h4(e13)) | (e23 = h4(e12)) | (e23 = h4(e11)) | (e23 = h4(e10))) & ((e22 = h4(e13)) | (e22 = h4(e12)) | (e22 = h4(e11)) | (e22 = h4(e10))) & ((e21 = h4(e13)) | (e21 = h4(e12)) | (e21 = h4(e11)) | (e21 = h4(e10))) & ((e20 = h4(e13)) | (e20 = h4(e12)) | (e20 = h4(e11)) | (e20 = h4(e10))) & (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) & (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) & (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) & (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) & (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) & (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) & (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) & (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) & (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) & (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) & (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) & (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) & (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) & (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) & (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) & (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) | (((e23 = h3(e13)) | (e23 = h3(e12)) | (e23 = h3(e11)) | (e23 = h3(e10))) & ((e22 = h3(e13)) | (e22 = h3(e12)) | (e22 = h3(e11)) | (e22 = h3(e10))) & ((e21 = h3(e13)) | (e21 = h3(e12)) | (e21 = h3(e11)) | (e21 = h3(e10))) & ((e20 = h3(e13)) | (e20 = h3(e12)) | (e20 = h3(e11)) | (e20 = h3(e10))) & (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) & (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) & (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) & (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) & (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) & (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) & (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) & (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) & (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) & (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) & (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) & (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) & (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) & (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) & (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) & (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) | (((e23 = h2(e13)) | (e23 = h2(e12)) | (e23 = h2(e11)) | (e23 = h2(e10))) & ((e22 = h2(e13)) | (e22 = h2(e12)) | (e22 = h2(e11)) | (e22 = h2(e10))) & ((e21 = h2(e13)) | (e21 = h2(e12)) | (e21 = h2(e11)) | (e21 = h2(e10))) & ((e20 = h2(e13)) | (e20 = h2(e12)) | (e20 = h2(e11)) | (e20 = h2(e10))) & (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) & (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) & (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) & (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) & (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) & (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) & (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) & (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) & (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) & (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) & (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) & (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) & (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) & (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) & (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) & (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) | (((e23 = h1(e13)) | (e23 = h1(e12)) | (e23 = h1(e11)) | (e23 = h1(e10))) & ((e22 = h1(e13)) | (e22 = h1(e12)) | (e22 = h1(e11)) | (e22 = h1(e10))) & ((e21 = h1(e13)) | (e21 = h1(e12)) | (e21 = h1(e11)) | (e21 = h1(e10))) & ((e20 = h1(e13)) | (e20 = h1(e12)) | (e20 = h1(e11)) | (e20 = h1(e10))) & (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) & (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) & (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) & (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) & (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) & (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) & (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) & (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) & (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) & (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) & (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) & (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) & (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) & (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) & (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) & (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG115+1.p', co1)).
fof(f3055, plain, ~ spl144_295, inference(avatar_split_clause, [], [f3054, f3046])).
fof(f3054, plain, ~ sP141, inference(subsumption_resolution, [], [f1114, f3028])).
fof(f1114, plain, (~ (e20 = h4(e10)) | ~ sP141), inference(cnf_transformation, [], [f302])).
fof(f302, plain, ((~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10))) | ~ sP141), inference(nnf_transformation, [], [f164])).
fof(f3041, plain, (~ spl144_293 | ~ spl144_294), inference(avatar_split_clause, [], [f1111, f3038, f3034])).
fof(f1111, plain, (~ (e21 = h4(e11)) | ~ sP142), inference(cnf_transformation, [], [f301])).
fof(f301, plain, ((~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10))) | ~ sP142), inference(nnf_transformation, [], [f165])).
fof(f2988, plain, spl144_66, inference(avatar_split_clause, [], [f1088, f1479])).
fof(f1088, plain, (e21 = op2(e23, e23)), inference(cnf_transformation, [], [f13])).
fof(f2986, plain, spl144_2, inference(avatar_split_clause, [], [f1085, f1175])).
fof(f1085, plain, (e11 = op1(e13, e13)), inference(cnf_transformation, [], [f12])).
fof(f2985, plain, (spl144_125 | spl144_105 | spl144_85 | spl144_65), inference(avatar_split_clause, [], [f1072, f1475, f1560, f1645, f1730])).
fof(f1072, plain, ((e20 = op2(e23, e23)) | (e20 = op2(e22, e22)) | (e20 = op2(e21, e21)) | (e20 = op2(e20, e20))), inference(cnf_transformation, [], [f154])).
fof(f154, plain, ((((e23 = op2(e23, e23)) & (e23 = op2(e23, e23)) & (e23 = op2(e23, e23)) & (e23 = op2(e23, e23))) | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69) & ((~ (e23 = op2(e23, op2(e23, e23))) & ~ (e22 = op2(e23, op2(e23, e22))) & ~ (e21 = op2(e23, op2(e23, e21))) & ~ (e20 = op2(e23, op2(e23, e20)))) | sP68 | sP67 | sP66) & ((e23 = op2(e23, e23)) | (e23 = op2(e22, e22)) | (e23 = op2(e21, e21)) | (op2(e20, e20) = e23)) & ((e22 = op2(e23, e23)) | (e22 = op2(e22, e22)) | (e22 = op2(e21, e21)) | (op2(e20, e20) = e22)) & ((e21 = op2(e23, e23)) | (e21 = op2(e22, e22)) | (e21 = op2(e21, e21)) | (op2(e20, e20) = e21)) & ((e20 = op2(e23, e23)) | (e20 = op2(e22, e22)) | (e20 = op2(e21, e21)) | (e20 = op2(e20, e20)))), inference(definition_folding, [], [f11, e153, e152, e151, e150, e149, e148, e147, e146, e145, e144, e143, e142, e141, e140, e139, e138, e137, e136, e135, e134, e133, e132, e131, e130, e129, e128, e127, e126, e125, e124, e123, e122, e121, e120, e119, e118, e117, e116, e115, e114, e113, e112, e111, e110, e109, e108, e107, e106, e105, e104, e103, e102, e101, e100, e99, e98, e97, e96, e95, e94, e93, e92, e91, e90, e89, e88])).
fof(f88, plain, ((~ (e23 = op2(e20, op2(e20, e23))) & ~ (e22 = op2(e20, op2(e20, e22))) & ~ (e21 = op2(e20, op2(e20, e21))) & ~ (e20 = op2(e20, op2(e20, e20)))) | ~ sP66), inference(usedef, [], [e88])).
fof(e88, plain, (sP66 <=> (~ (e23 = op2(e20, op2(e20, e23))) & ~ (e22 = op2(e20, op2(e20, e22))) & ~ (e21 = op2(e20, op2(e20, e21))) & ~ (e20 = op2(e20, op2(e20, e20))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP66])])).
fof(f89, plain, ((~ (e23 = op2(e21, op2(e21, e23))) & ~ (e22 = op2(e21, op2(e21, e22))) & ~ (e21 = op2(e21, op2(e21, e21))) & ~ (e20 = op2(e21, op2(e21, e20)))) | ~ sP67), inference(usedef, [], [e89])).
fof(e89, plain, (sP67 <=> (~ (e23 = op2(e21, op2(e21, e23))) & ~ (e22 = op2(e21, op2(e21, e22))) & ~ (e21 = op2(e21, op2(e21, e21))) & ~ (e20 = op2(e21, op2(e21, e20))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP67])])).
fof(f90, plain, ((~ (e23 = op2(e22, op2(e22, e23))) & ~ (e22 = op2(e22, op2(e22, e22))) & ~ (e21 = op2(e22, op2(e22, e21))) & ~ (e20 = op2(e22, op2(e22, e20)))) | ~ sP68), inference(usedef, [], [e90])).
fof(e90, plain, (sP68 <=> (~ (e23 = op2(e22, op2(e22, e23))) & ~ (e22 = op2(e22, op2(e22, e22))) & ~ (e21 = op2(e22, op2(e22, e21))) & ~ (e20 = op2(e22, op2(e22, e20))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP68])])).
fof(f91, plain, (((e20 = op2(e20, e20)) & (e20 = op2(e20, e20)) & (e20 = op2(e20, e20)) & (e20 = op2(e20, e20))) | ~ sP69), inference(usedef, [], [e91])).
fof(e91, plain, (sP69 <=> ((e20 = op2(e20, e20)) & (e20 = op2(e20, e20)) & (e20 = op2(e20, e20)) & (e20 = op2(e20, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP69])])).
fof(f92, plain, (((e20 = op2(e20, e21)) & (e20 = op2(e21, e20)) & (e20 = op2(e20, e20)) & (e20 = op2(e20, e20))) | ~ sP70), inference(usedef, [], [e92])).
fof(e92, plain, (sP70 <=> ((e20 = op2(e20, e21)) & (e20 = op2(e21, e20)) & (e20 = op2(e20, e20)) & (e20 = op2(e20, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP70])])).
fof(f93, plain, (((e20 = op2(e20, e22)) & (e20 = op2(e22, e20)) & (e20 = op2(e20, e20)) & (e20 = op2(e20, e20))) | ~ sP71), inference(usedef, [], [e93])).
fof(e93, plain, (sP71 <=> ((e20 = op2(e20, e22)) & (e20 = op2(e22, e20)) & (e20 = op2(e20, e20)) & (e20 = op2(e20, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP71])])).
fof(f94, plain, (((e20 = op2(e20, e23)) & (e20 = op2(e23, e20)) & (e20 = op2(e20, e20)) & (e20 = op2(e20, e20))) | ~ sP72), inference(usedef, [], [e94])).
fof(e94, plain, (sP72 <=> ((e20 = op2(e20, e23)) & (e20 = op2(e23, e20)) & (e20 = op2(e20, e20)) & (e20 = op2(e20, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP72])])).
fof(f95, plain, (((e21 = op2(e21, e20)) & (e21 = op2(e20, e21)) & (e20 = op2(e21, e20)) & (e20 = op2(e20, e21))) | ~ sP73), inference(usedef, [], [e95])).
fof(e95, plain, (sP73 <=> ((e21 = op2(e21, e20)) & (e21 = op2(e20, e21)) & (e20 = op2(e21, e20)) & (e20 = op2(e20, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP73])])).
fof(f96, plain, (((e21 = op2(e21, e21)) & (e21 = op2(e21, e21)) & (e20 = op2(e21, e20)) & (e20 = op2(e20, e21))) | ~ sP74), inference(usedef, [], [e96])).
fof(e96, plain, (sP74 <=> ((e21 = op2(e21, e21)) & (e21 = op2(e21, e21)) & (e20 = op2(e21, e20)) & (e20 = op2(e20, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP74])])).
fof(f97, plain, (((e21 = op2(e21, e22)) & (e21 = op2(e22, e21)) & (e20 = op2(e21, e20)) & (e20 = op2(e20, e21))) | ~ sP75), inference(usedef, [], [e97])).
fof(e97, plain, (sP75 <=> ((e21 = op2(e21, e22)) & (e21 = op2(e22, e21)) & (e20 = op2(e21, e20)) & (e20 = op2(e20, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP75])])).
fof(f98, plain, (((e21 = op2(e21, e23)) & (e21 = op2(e23, e21)) & (e20 = op2(e21, e20)) & (e20 = op2(e20, e21))) | ~ sP76), inference(usedef, [], [e98])).
fof(e98, plain, (sP76 <=> ((e21 = op2(e21, e23)) & (e21 = op2(e23, e21)) & (e20 = op2(e21, e20)) & (e20 = op2(e20, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP76])])).
fof(f99, plain, (((e22 = op2(e22, e20)) & (e22 = op2(e20, e22)) & (e20 = op2(e22, e20)) & (e20 = op2(e20, e22))) | ~ sP77), inference(usedef, [], [e99])).
fof(e99, plain, (sP77 <=> ((e22 = op2(e22, e20)) & (e22 = op2(e20, e22)) & (e20 = op2(e22, e20)) & (e20 = op2(e20, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP77])])).
fof(f100, plain, (((e22 = op2(e22, e21)) & (e22 = op2(e21, e22)) & (e20 = op2(e22, e20)) & (e20 = op2(e20, e22))) | ~ sP78), inference(usedef, [], [e100])).
fof(e100, plain, (sP78 <=> ((e22 = op2(e22, e21)) & (e22 = op2(e21, e22)) & (e20 = op2(e22, e20)) & (e20 = op2(e20, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP78])])).
fof(f101, plain, (((e22 = op2(e22, e22)) & (e22 = op2(e22, e22)) & (e20 = op2(e22, e20)) & (e20 = op2(e20, e22))) | ~ sP79), inference(usedef, [], [e101])).
fof(e101, plain, (sP79 <=> ((e22 = op2(e22, e22)) & (e22 = op2(e22, e22)) & (e20 = op2(e22, e20)) & (e20 = op2(e20, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP79])])).
fof(f102, plain, (((e22 = op2(e22, e23)) & (e22 = op2(e23, e22)) & (e20 = op2(e22, e20)) & (e20 = op2(e20, e22))) | ~ sP80), inference(usedef, [], [e102])).
fof(e102, plain, (sP80 <=> ((e22 = op2(e22, e23)) & (e22 = op2(e23, e22)) & (e20 = op2(e22, e20)) & (e20 = op2(e20, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP80])])).
fof(f103, plain, (((e23 = op2(e23, e20)) & (e23 = op2(e20, e23)) & (e20 = op2(e23, e20)) & (e20 = op2(e20, e23))) | ~ sP81), inference(usedef, [], [e103])).
fof(e103, plain, (sP81 <=> ((e23 = op2(e23, e20)) & (e23 = op2(e20, e23)) & (e20 = op2(e23, e20)) & (e20 = op2(e20, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP81])])).
fof(f104, plain, (((e23 = op2(e23, e21)) & (e23 = op2(e21, e23)) & (e20 = op2(e23, e20)) & (e20 = op2(e20, e23))) | ~ sP82), inference(usedef, [], [e104])).
fof(e104, plain, (sP82 <=> ((e23 = op2(e23, e21)) & (e23 = op2(e21, e23)) & (e20 = op2(e23, e20)) & (e20 = op2(e20, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP82])])).
fof(f105, plain, (((e23 = op2(e23, e22)) & (e23 = op2(e22, e23)) & (e20 = op2(e23, e20)) & (e20 = op2(e20, e23))) | ~ sP83), inference(usedef, [], [e105])).
fof(e105, plain, (sP83 <=> ((e23 = op2(e23, e22)) & (e23 = op2(e22, e23)) & (e20 = op2(e23, e20)) & (e20 = op2(e20, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP83])])).
fof(f106, plain, (((e23 = op2(e23, e23)) & (e23 = op2(e23, e23)) & (e20 = op2(e23, e20)) & (e20 = op2(e20, e23))) | ~ sP84), inference(usedef, [], [e106])).
fof(e106, plain, (sP84 <=> ((e23 = op2(e23, e23)) & (e23 = op2(e23, e23)) & (e20 = op2(e23, e20)) & (e20 = op2(e20, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP84])])).
fof(f107, plain, (((e20 = op2(e20, e20)) & (e20 = op2(e20, e20)) & (e21 = op2(e20, e21)) & (e21 = op2(e21, e20))) | ~ sP85), inference(usedef, [], [e107])).
fof(e107, plain, (sP85 <=> ((e20 = op2(e20, e20)) & (e20 = op2(e20, e20)) & (e21 = op2(e20, e21)) & (e21 = op2(e21, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP85])])).
fof(f108, plain, (((e20 = op2(e20, e21)) & (e20 = op2(e21, e20)) & (e21 = op2(e20, e21)) & (e21 = op2(e21, e20))) | ~ sP86), inference(usedef, [], [e108])).
fof(e108, plain, (sP86 <=> ((e20 = op2(e20, e21)) & (e20 = op2(e21, e20)) & (e21 = op2(e20, e21)) & (e21 = op2(e21, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP86])])).
fof(f109, plain, (((e20 = op2(e20, e22)) & (e20 = op2(e22, e20)) & (e21 = op2(e20, e21)) & (e21 = op2(e21, e20))) | ~ sP87), inference(usedef, [], [e109])).
fof(e109, plain, (sP87 <=> ((e20 = op2(e20, e22)) & (e20 = op2(e22, e20)) & (e21 = op2(e20, e21)) & (e21 = op2(e21, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP87])])).
fof(f110, plain, (((e20 = op2(e20, e23)) & (e20 = op2(e23, e20)) & (e21 = op2(e20, e21)) & (e21 = op2(e21, e20))) | ~ sP88), inference(usedef, [], [e110])).
fof(e110, plain, (sP88 <=> ((e20 = op2(e20, e23)) & (e20 = op2(e23, e20)) & (e21 = op2(e20, e21)) & (e21 = op2(e21, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP88])])).
fof(f111, plain, (((e21 = op2(e21, e20)) & (e21 = op2(e20, e21)) & (e21 = op2(e21, e21)) & (e21 = op2(e21, e21))) | ~ sP89), inference(usedef, [], [e111])).
fof(e111, plain, (sP89 <=> ((e21 = op2(e21, e20)) & (e21 = op2(e20, e21)) & (e21 = op2(e21, e21)) & (e21 = op2(e21, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP89])])).
fof(f112, plain, (((e21 = op2(e21, e21)) & (e21 = op2(e21, e21)) & (e21 = op2(e21, e21)) & (e21 = op2(e21, e21))) | ~ sP90), inference(usedef, [], [e112])).
fof(e112, plain, (sP90 <=> ((e21 = op2(e21, e21)) & (e21 = op2(e21, e21)) & (e21 = op2(e21, e21)) & (e21 = op2(e21, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP90])])).
fof(f113, plain, (((e21 = op2(e21, e22)) & (e21 = op2(e22, e21)) & (e21 = op2(e21, e21)) & (e21 = op2(e21, e21))) | ~ sP91), inference(usedef, [], [e113])).
fof(e113, plain, (sP91 <=> ((e21 = op2(e21, e22)) & (e21 = op2(e22, e21)) & (e21 = op2(e21, e21)) & (e21 = op2(e21, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP91])])).
fof(f114, plain, (((e21 = op2(e21, e23)) & (e21 = op2(e23, e21)) & (e21 = op2(e21, e21)) & (e21 = op2(e21, e21))) | ~ sP92), inference(usedef, [], [e114])).
fof(e114, plain, (sP92 <=> ((e21 = op2(e21, e23)) & (e21 = op2(e23, e21)) & (e21 = op2(e21, e21)) & (e21 = op2(e21, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP92])])).
fof(f115, plain, (((e22 = op2(e22, e20)) & (e22 = op2(e20, e22)) & (e21 = op2(e22, e21)) & (e21 = op2(e21, e22))) | ~ sP93), inference(usedef, [], [e115])).
fof(e115, plain, (sP93 <=> ((e22 = op2(e22, e20)) & (e22 = op2(e20, e22)) & (e21 = op2(e22, e21)) & (e21 = op2(e21, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP93])])).
fof(f116, plain, (((e22 = op2(e22, e21)) & (e22 = op2(e21, e22)) & (e21 = op2(e22, e21)) & (e21 = op2(e21, e22))) | ~ sP94), inference(usedef, [], [e116])).
fof(e116, plain, (sP94 <=> ((e22 = op2(e22, e21)) & (e22 = op2(e21, e22)) & (e21 = op2(e22, e21)) & (e21 = op2(e21, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP94])])).
fof(f117, plain, (((e22 = op2(e22, e22)) & (e22 = op2(e22, e22)) & (e21 = op2(e22, e21)) & (e21 = op2(e21, e22))) | ~ sP95), inference(usedef, [], [e117])).
fof(e117, plain, (sP95 <=> ((e22 = op2(e22, e22)) & (e22 = op2(e22, e22)) & (e21 = op2(e22, e21)) & (e21 = op2(e21, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP95])])).
fof(f118, plain, (((e22 = op2(e22, e23)) & (e22 = op2(e23, e22)) & (e21 = op2(e22, e21)) & (e21 = op2(e21, e22))) | ~ sP96), inference(usedef, [], [e118])).
fof(e118, plain, (sP96 <=> ((e22 = op2(e22, e23)) & (e22 = op2(e23, e22)) & (e21 = op2(e22, e21)) & (e21 = op2(e21, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP96])])).
fof(f119, plain, (((e23 = op2(e23, e20)) & (e23 = op2(e20, e23)) & (e21 = op2(e23, e21)) & (e21 = op2(e21, e23))) | ~ sP97), inference(usedef, [], [e119])).
fof(e119, plain, (sP97 <=> ((e23 = op2(e23, e20)) & (e23 = op2(e20, e23)) & (e21 = op2(e23, e21)) & (e21 = op2(e21, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP97])])).
fof(f120, plain, (((e23 = op2(e23, e21)) & (e23 = op2(e21, e23)) & (e21 = op2(e23, e21)) & (e21 = op2(e21, e23))) | ~ sP98), inference(usedef, [], [e120])).
fof(e120, plain, (sP98 <=> ((e23 = op2(e23, e21)) & (e23 = op2(e21, e23)) & (e21 = op2(e23, e21)) & (e21 = op2(e21, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP98])])).
fof(f121, plain, (((e23 = op2(e23, e22)) & (e23 = op2(e22, e23)) & (e21 = op2(e23, e21)) & (e21 = op2(e21, e23))) | ~ sP99), inference(usedef, [], [e121])).
fof(e121, plain, (sP99 <=> ((e23 = op2(e23, e22)) & (e23 = op2(e22, e23)) & (e21 = op2(e23, e21)) & (e21 = op2(e21, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP99])])).
fof(f122, plain, (((e23 = op2(e23, e23)) & (e23 = op2(e23, e23)) & (e21 = op2(e23, e21)) & (e21 = op2(e21, e23))) | ~ sP100), inference(usedef, [], [e122])).
fof(e122, plain, (sP100 <=> ((e23 = op2(e23, e23)) & (e23 = op2(e23, e23)) & (e21 = op2(e23, e21)) & (e21 = op2(e21, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP100])])).
fof(f123, plain, (((e20 = op2(e20, e20)) & (e20 = op2(e20, e20)) & (e22 = op2(e20, e22)) & (e22 = op2(e22, e20))) | ~ sP101), inference(usedef, [], [e123])).
fof(e123, plain, (sP101 <=> ((e20 = op2(e20, e20)) & (e20 = op2(e20, e20)) & (e22 = op2(e20, e22)) & (e22 = op2(e22, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP101])])).
fof(f124, plain, (((e20 = op2(e20, e21)) & (e20 = op2(e21, e20)) & (e22 = op2(e20, e22)) & (e22 = op2(e22, e20))) | ~ sP102), inference(usedef, [], [e124])).
fof(e124, plain, (sP102 <=> ((e20 = op2(e20, e21)) & (e20 = op2(e21, e20)) & (e22 = op2(e20, e22)) & (e22 = op2(e22, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP102])])).
fof(f125, plain, (((e20 = op2(e20, e22)) & (e20 = op2(e22, e20)) & (e22 = op2(e20, e22)) & (e22 = op2(e22, e20))) | ~ sP103), inference(usedef, [], [e125])).
fof(e125, plain, (sP103 <=> ((e20 = op2(e20, e22)) & (e20 = op2(e22, e20)) & (e22 = op2(e20, e22)) & (e22 = op2(e22, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP103])])).
fof(f126, plain, (((e20 = op2(e20, e23)) & (e20 = op2(e23, e20)) & (e22 = op2(e20, e22)) & (e22 = op2(e22, e20))) | ~ sP104), inference(usedef, [], [e126])).
fof(e126, plain, (sP104 <=> ((e20 = op2(e20, e23)) & (e20 = op2(e23, e20)) & (e22 = op2(e20, e22)) & (e22 = op2(e22, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP104])])).
fof(f127, plain, (((e21 = op2(e21, e20)) & (e21 = op2(e20, e21)) & (e22 = op2(e21, e22)) & (e22 = op2(e22, e21))) | ~ sP105), inference(usedef, [], [e127])).
fof(e127, plain, (sP105 <=> ((e21 = op2(e21, e20)) & (e21 = op2(e20, e21)) & (e22 = op2(e21, e22)) & (e22 = op2(e22, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP105])])).
fof(f128, plain, (((e21 = op2(e21, e21)) & (e21 = op2(e21, e21)) & (e22 = op2(e21, e22)) & (e22 = op2(e22, e21))) | ~ sP106), inference(usedef, [], [e128])).
fof(e128, plain, (sP106 <=> ((e21 = op2(e21, e21)) & (e21 = op2(e21, e21)) & (e22 = op2(e21, e22)) & (e22 = op2(e22, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP106])])).
fof(f129, plain, (((e21 = op2(e21, e22)) & (e21 = op2(e22, e21)) & (e22 = op2(e21, e22)) & (e22 = op2(e22, e21))) | ~ sP107), inference(usedef, [], [e129])).
fof(e129, plain, (sP107 <=> ((e21 = op2(e21, e22)) & (e21 = op2(e22, e21)) & (e22 = op2(e21, e22)) & (e22 = op2(e22, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP107])])).
fof(f130, plain, (((e21 = op2(e21, e23)) & (e21 = op2(e23, e21)) & (e22 = op2(e21, e22)) & (e22 = op2(e22, e21))) | ~ sP108), inference(usedef, [], [e130])).
fof(e130, plain, (sP108 <=> ((e21 = op2(e21, e23)) & (e21 = op2(e23, e21)) & (e22 = op2(e21, e22)) & (e22 = op2(e22, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP108])])).
fof(f131, plain, (((e22 = op2(e22, e20)) & (e22 = op2(e20, e22)) & (e22 = op2(e22, e22)) & (e22 = op2(e22, e22))) | ~ sP109), inference(usedef, [], [e131])).
fof(e131, plain, (sP109 <=> ((e22 = op2(e22, e20)) & (e22 = op2(e20, e22)) & (e22 = op2(e22, e22)) & (e22 = op2(e22, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP109])])).
fof(f132, plain, (((e22 = op2(e22, e21)) & (e22 = op2(e21, e22)) & (e22 = op2(e22, e22)) & (e22 = op2(e22, e22))) | ~ sP110), inference(usedef, [], [e132])).
fof(e132, plain, (sP110 <=> ((e22 = op2(e22, e21)) & (e22 = op2(e21, e22)) & (e22 = op2(e22, e22)) & (e22 = op2(e22, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP110])])).
fof(f133, plain, (((e22 = op2(e22, e22)) & (e22 = op2(e22, e22)) & (e22 = op2(e22, e22)) & (e22 = op2(e22, e22))) | ~ sP111), inference(usedef, [], [e133])).
fof(e133, plain, (sP111 <=> ((e22 = op2(e22, e22)) & (e22 = op2(e22, e22)) & (e22 = op2(e22, e22)) & (e22 = op2(e22, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP111])])).
fof(f134, plain, (((e22 = op2(e22, e23)) & (e22 = op2(e23, e22)) & (e22 = op2(e22, e22)) & (e22 = op2(e22, e22))) | ~ sP112), inference(usedef, [], [e134])).
fof(e134, plain, (sP112 <=> ((e22 = op2(e22, e23)) & (e22 = op2(e23, e22)) & (e22 = op2(e22, e22)) & (e22 = op2(e22, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP112])])).
fof(f135, plain, (((e23 = op2(e23, e20)) & (e23 = op2(e20, e23)) & (e22 = op2(e23, e22)) & (e22 = op2(e22, e23))) | ~ sP113), inference(usedef, [], [e135])).
fof(e135, plain, (sP113 <=> ((e23 = op2(e23, e20)) & (e23 = op2(e20, e23)) & (e22 = op2(e23, e22)) & (e22 = op2(e22, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP113])])).
fof(f136, plain, (((e23 = op2(e23, e21)) & (e23 = op2(e21, e23)) & (e22 = op2(e23, e22)) & (e22 = op2(e22, e23))) | ~ sP114), inference(usedef, [], [e136])).
fof(e136, plain, (sP114 <=> ((e23 = op2(e23, e21)) & (e23 = op2(e21, e23)) & (e22 = op2(e23, e22)) & (e22 = op2(e22, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP114])])).
fof(f137, plain, (((e23 = op2(e23, e22)) & (e23 = op2(e22, e23)) & (e22 = op2(e23, e22)) & (e22 = op2(e22, e23))) | ~ sP115), inference(usedef, [], [e137])).
fof(e137, plain, (sP115 <=> ((e23 = op2(e23, e22)) & (e23 = op2(e22, e23)) & (e22 = op2(e23, e22)) & (e22 = op2(e22, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP115])])).
fof(f138, plain, (((e23 = op2(e23, e23)) & (e23 = op2(e23, e23)) & (e22 = op2(e23, e22)) & (e22 = op2(e22, e23))) | ~ sP116), inference(usedef, [], [e138])).
fof(e138, plain, (sP116 <=> ((e23 = op2(e23, e23)) & (e23 = op2(e23, e23)) & (e22 = op2(e23, e22)) & (e22 = op2(e22, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP116])])).
fof(f139, plain, (((e20 = op2(e20, e20)) & (e20 = op2(e20, e20)) & (e23 = op2(e20, e23)) & (e23 = op2(e23, e20))) | ~ sP117), inference(usedef, [], [e139])).
fof(e139, plain, (sP117 <=> ((e20 = op2(e20, e20)) & (e20 = op2(e20, e20)) & (e23 = op2(e20, e23)) & (e23 = op2(e23, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP117])])).
fof(f140, plain, (((e20 = op2(e20, e21)) & (e20 = op2(e21, e20)) & (e23 = op2(e20, e23)) & (e23 = op2(e23, e20))) | ~ sP118), inference(usedef, [], [e140])).
fof(e140, plain, (sP118 <=> ((e20 = op2(e20, e21)) & (e20 = op2(e21, e20)) & (e23 = op2(e20, e23)) & (e23 = op2(e23, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP118])])).
fof(f141, plain, (((e20 = op2(e20, e22)) & (e20 = op2(e22, e20)) & (e23 = op2(e20, e23)) & (e23 = op2(e23, e20))) | ~ sP119), inference(usedef, [], [e141])).
fof(e141, plain, (sP119 <=> ((e20 = op2(e20, e22)) & (e20 = op2(e22, e20)) & (e23 = op2(e20, e23)) & (e23 = op2(e23, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP119])])).
fof(f142, plain, (((e20 = op2(e20, e23)) & (e20 = op2(e23, e20)) & (e23 = op2(e20, e23)) & (e23 = op2(e23, e20))) | ~ sP120), inference(usedef, [], [e142])).
fof(e142, plain, (sP120 <=> ((e20 = op2(e20, e23)) & (e20 = op2(e23, e20)) & (e23 = op2(e20, e23)) & (e23 = op2(e23, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP120])])).
fof(f143, plain, (((e21 = op2(e21, e20)) & (e21 = op2(e20, e21)) & (e23 = op2(e21, e23)) & (e23 = op2(e23, e21))) | ~ sP121), inference(usedef, [], [e143])).
fof(e143, plain, (sP121 <=> ((e21 = op2(e21, e20)) & (e21 = op2(e20, e21)) & (e23 = op2(e21, e23)) & (e23 = op2(e23, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP121])])).
fof(f144, plain, (((e21 = op2(e21, e21)) & (e21 = op2(e21, e21)) & (e23 = op2(e21, e23)) & (e23 = op2(e23, e21))) | ~ sP122), inference(usedef, [], [e144])).
fof(e144, plain, (sP122 <=> ((e21 = op2(e21, e21)) & (e21 = op2(e21, e21)) & (e23 = op2(e21, e23)) & (e23 = op2(e23, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP122])])).
fof(f145, plain, (((e21 = op2(e21, e22)) & (e21 = op2(e22, e21)) & (e23 = op2(e21, e23)) & (e23 = op2(e23, e21))) | ~ sP123), inference(usedef, [], [e145])).
fof(e145, plain, (sP123 <=> ((e21 = op2(e21, e22)) & (e21 = op2(e22, e21)) & (e23 = op2(e21, e23)) & (e23 = op2(e23, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP123])])).
fof(f146, plain, (((e21 = op2(e21, e23)) & (e21 = op2(e23, e21)) & (e23 = op2(e21, e23)) & (e23 = op2(e23, e21))) | ~ sP124), inference(usedef, [], [e146])).
fof(e146, plain, (sP124 <=> ((e21 = op2(e21, e23)) & (e21 = op2(e23, e21)) & (e23 = op2(e21, e23)) & (e23 = op2(e23, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP124])])).
fof(f147, plain, (((e22 = op2(e22, e20)) & (e22 = op2(e20, e22)) & (e23 = op2(e22, e23)) & (e23 = op2(e23, e22))) | ~ sP125), inference(usedef, [], [e147])).
fof(e147, plain, (sP125 <=> ((e22 = op2(e22, e20)) & (e22 = op2(e20, e22)) & (e23 = op2(e22, e23)) & (e23 = op2(e23, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP125])])).
fof(f148, plain, (((e22 = op2(e22, e21)) & (e22 = op2(e21, e22)) & (e23 = op2(e22, e23)) & (e23 = op2(e23, e22))) | ~ sP126), inference(usedef, [], [e148])).
fof(e148, plain, (sP126 <=> ((e22 = op2(e22, e21)) & (e22 = op2(e21, e22)) & (e23 = op2(e22, e23)) & (e23 = op2(e23, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP126])])).
fof(f149, plain, (((e22 = op2(e22, e22)) & (e22 = op2(e22, e22)) & (e23 = op2(e22, e23)) & (e23 = op2(e23, e22))) | ~ sP127), inference(usedef, [], [e149])).
fof(e149, plain, (sP127 <=> ((e22 = op2(e22, e22)) & (e22 = op2(e22, e22)) & (e23 = op2(e22, e23)) & (e23 = op2(e23, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP127])])).
fof(f150, plain, (((e22 = op2(e22, e23)) & (e22 = op2(e23, e22)) & (e23 = op2(e22, e23)) & (e23 = op2(e23, e22))) | ~ sP128), inference(usedef, [], [e150])).
fof(e150, plain, (sP128 <=> ((e22 = op2(e22, e23)) & (e22 = op2(e23, e22)) & (e23 = op2(e22, e23)) & (e23 = op2(e23, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP128])])).
fof(f151, plain, (((e23 = op2(e23, e20)) & (e23 = op2(e20, e23)) & (e23 = op2(e23, e23)) & (e23 = op2(e23, e23))) | ~ sP129), inference(usedef, [], [e151])).
fof(e151, plain, (sP129 <=> ((e23 = op2(e23, e20)) & (e23 = op2(e20, e23)) & (e23 = op2(e23, e23)) & (e23 = op2(e23, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP129])])).
fof(f152, plain, (((e23 = op2(e23, e21)) & (e23 = op2(e21, e23)) & (e23 = op2(e23, e23)) & (e23 = op2(e23, e23))) | ~ sP130), inference(usedef, [], [e152])).
fof(e152, plain, (sP130 <=> ((e23 = op2(e23, e21)) & (e23 = op2(e21, e23)) & (e23 = op2(e23, e23)) & (e23 = op2(e23, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP130])])).
fof(f153, plain, (((e23 = op2(e23, e22)) & (e23 = op2(e22, e23)) & (e23 = op2(e23, e23)) & (e23 = op2(e23, e23))) | ~ sP131), inference(usedef, [], [e153])).
fof(e153, plain, (sP131 <=> ((e23 = op2(e23, e22)) & (e23 = op2(e22, e23)) & (e23 = op2(e23, e23)) & (e23 = op2(e23, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP131])])).
fof(f11, plain, ((((e23 = op2(e23, e23)) & (e23 = op2(e23, e23)) & (e23 = op2(e23, e23)) & (e23 = op2(e23, e23))) | ((e23 = op2(e23, e22)) & (e23 = op2(e22, e23)) & (e23 = op2(e23, e23)) & (e23 = op2(e23, e23))) | ((e23 = op2(e23, e21)) & (e23 = op2(e21, e23)) & (e23 = op2(e23, e23)) & (e23 = op2(e23, e23))) | ((e23 = op2(e23, e20)) & (e23 = op2(e20, e23)) & (e23 = op2(e23, e23)) & (e23 = op2(e23, e23))) | ((e22 = op2(e22, e23)) & (e22 = op2(e23, e22)) & (e23 = op2(e22, e23)) & (e23 = op2(e23, e22))) | ((e22 = op2(e22, e22)) & (e22 = op2(e22, e22)) & (e23 = op2(e22, e23)) & (e23 = op2(e23, e22))) | ((e22 = op2(e22, e21)) & (e22 = op2(e21, e22)) & (e23 = op2(e22, e23)) & (e23 = op2(e23, e22))) | ((e22 = op2(e22, e20)) & (e22 = op2(e20, e22)) & (e23 = op2(e22, e23)) & (e23 = op2(e23, e22))) | ((e21 = op2(e21, e23)) & (e21 = op2(e23, e21)) & (e23 = op2(e21, e23)) & (e23 = op2(e23, e21))) | ((e21 = op2(e21, e22)) & (e21 = op2(e22, e21)) & (e23 = op2(e21, e23)) & (e23 = op2(e23, e21))) | ((e21 = op2(e21, e21)) & (e21 = op2(e21, e21)) & (e23 = op2(e21, e23)) & (e23 = op2(e23, e21))) | ((e21 = op2(e21, e20)) & (e21 = op2(e20, e21)) & (e23 = op2(e21, e23)) & (e23 = op2(e23, e21))) | ((e20 = op2(e20, e23)) & (e20 = op2(e23, e20)) & (e23 = op2(e20, e23)) & (e23 = op2(e23, e20))) | ((e20 = op2(e20, e22)) & (e20 = op2(e22, e20)) & (e23 = op2(e20, e23)) & (e23 = op2(e23, e20))) | ((e20 = op2(e20, e21)) & (e20 = op2(e21, e20)) & (e23 = op2(e20, e23)) & (e23 = op2(e23, e20))) | ((e20 = op2(e20, e20)) & (e20 = op2(e20, e20)) & (e23 = op2(e20, e23)) & (e23 = op2(e23, e20))) | ((e23 = op2(e23, e23)) & (e23 = op2(e23, e23)) & (e22 = op2(e23, e22)) & (e22 = op2(e22, e23))) | ((e23 = op2(e23, e22)) & (e23 = op2(e22, e23)) & (e22 = op2(e23, e22)) & (e22 = op2(e22, e23))) | ((e23 = op2(e23, e21)) & (e23 = op2(e21, e23)) & (e22 = op2(e23, e22)) & (e22 = op2(e22, e23))) | ((e23 = op2(e23, e20)) & (e23 = op2(e20, e23)) & (e22 = op2(e23, e22)) & (e22 = op2(e22, e23))) | ((e22 = op2(e22, e23)) & (e22 = op2(e23, e22)) & (e22 = op2(e22, e22)) & (e22 = op2(e22, e22))) | ((e22 = op2(e22, e22)) & (e22 = op2(e22, e22)) & (e22 = op2(e22, e22)) & (e22 = op2(e22, e22))) | ((e22 = op2(e22, e21)) & (e22 = op2(e21, e22)) & (e22 = op2(e22, e22)) & (e22 = op2(e22, e22))) | ((e22 = op2(e22, e20)) & (e22 = op2(e20, e22)) & (e22 = op2(e22, e22)) & (e22 = op2(e22, e22))) | ((e21 = op2(e21, e23)) & (e21 = op2(e23, e21)) & (e22 = op2(e21, e22)) & (e22 = op2(e22, e21))) | ((e21 = op2(e21, e22)) & (e21 = op2(e22, e21)) & (e22 = op2(e21, e22)) & (e22 = op2(e22, e21))) | ((e21 = op2(e21, e21)) & (e21 = op2(e21, e21)) & (e22 = op2(e21, e22)) & (e22 = op2(e22, e21))) | ((e21 = op2(e21, e20)) & (e21 = op2(e20, e21)) & (e22 = op2(e21, e22)) & (e22 = op2(e22, e21))) | ((e20 = op2(e20, e23)) & (e20 = op2(e23, e20)) & (e22 = op2(e20, e22)) & (e22 = op2(e22, e20))) | ((e20 = op2(e20, e22)) & (e20 = op2(e22, e20)) & (e22 = op2(e20, e22)) & (e22 = op2(e22, e20))) | ((e20 = op2(e20, e21)) & (e20 = op2(e21, e20)) & (e22 = op2(e20, e22)) & (e22 = op2(e22, e20))) | ((e20 = op2(e20, e20)) & (e20 = op2(e20, e20)) & (e22 = op2(e20, e22)) & (e22 = op2(e22, e20))) | ((e23 = op2(e23, e23)) & (e23 = op2(e23, e23)) & (e21 = op2(e23, e21)) & (e21 = op2(e21, e23))) | ((e23 = op2(e23, e22)) & (e23 = op2(e22, e23)) & (e21 = op2(e23, e21)) & (e21 = op2(e21, e23))) | ((e23 = op2(e23, e21)) & (e23 = op2(e21, e23)) & (e21 = op2(e23, e21)) & (e21 = op2(e21, e23))) | ((e23 = op2(e23, e20)) & (e23 = op2(e20, e23)) & (e21 = op2(e23, e21)) & (e21 = op2(e21, e23))) | ((e22 = op2(e22, e23)) & (e22 = op2(e23, e22)) & (e21 = op2(e22, e21)) & (e21 = op2(e21, e22))) | ((e22 = op2(e22, e22)) & (e22 = op2(e22, e22)) & (e21 = op2(e22, e21)) & (e21 = op2(e21, e22))) | ((e22 = op2(e22, e21)) & (e22 = op2(e21, e22)) & (e21 = op2(e22, e21)) & (e21 = op2(e21, e22))) | ((e22 = op2(e22, e20)) & (e22 = op2(e20, e22)) & (e21 = op2(e22, e21)) & (e21 = op2(e21, e22))) | ((e21 = op2(e21, e23)) & (e21 = op2(e23, e21)) & (e21 = op2(e21, e21)) & (e21 = op2(e21, e21))) | ((e21 = op2(e21, e22)) & (e21 = op2(e22, e21)) & (e21 = op2(e21, e21)) & (e21 = op2(e21, e21))) | ((e21 = op2(e21, e21)) & (e21 = op2(e21, e21)) & (e21 = op2(e21, e21)) & (e21 = op2(e21, e21))) | ((e21 = op2(e21, e20)) & (e21 = op2(e20, e21)) & (e21 = op2(e21, e21)) & (e21 = op2(e21, e21))) | ((e20 = op2(e20, e23)) & (e20 = op2(e23, e20)) & (e21 = op2(e20, e21)) & (e21 = op2(e21, e20))) | ((e20 = op2(e20, e22)) & (e20 = op2(e22, e20)) & (e21 = op2(e20, e21)) & (e21 = op2(e21, e20))) | ((e20 = op2(e20, e21)) & (e20 = op2(e21, e20)) & (e21 = op2(e20, e21)) & (e21 = op2(e21, e20))) | ((e20 = op2(e20, e20)) & (e20 = op2(e20, e20)) & (e21 = op2(e20, e21)) & (e21 = op2(e21, e20))) | ((e23 = op2(e23, e23)) & (e23 = op2(e23, e23)) & (e20 = op2(e23, e20)) & (e20 = op2(e20, e23))) | ((e23 = op2(e23, e22)) & (e23 = op2(e22, e23)) & (e20 = op2(e23, e20)) & (e20 = op2(e20, e23))) | ((e23 = op2(e23, e21)) & (e23 = op2(e21, e23)) & (e20 = op2(e23, e20)) & (e20 = op2(e20, e23))) | ((e23 = op2(e23, e20)) & (e23 = op2(e20, e23)) & (e20 = op2(e23, e20)) & (e20 = op2(e20, e23))) | ((e22 = op2(e22, e23)) & (e22 = op2(e23, e22)) & (e20 = op2(e22, e20)) & (e20 = op2(e20, e22))) | ((e22 = op2(e22, e22)) & (e22 = op2(e22, e22)) & (e20 = op2(e22, e20)) & (e20 = op2(e20, e22))) | ((e22 = op2(e22, e21)) & (e22 = op2(e21, e22)) & (e20 = op2(e22, e20)) & (e20 = op2(e20, e22))) | ((e22 = op2(e22, e20)) & (e22 = op2(e20, e22)) & (e20 = op2(e22, e20)) & (e20 = op2(e20, e22))) | ((e21 = op2(e21, e23)) & (e21 = op2(e23, e21)) & (e20 = op2(e21, e20)) & (e20 = op2(e20, e21))) | ((e21 = op2(e21, e22)) & (e21 = op2(e22, e21)) & (e20 = op2(e21, e20)) & (e20 = op2(e20, e21))) | ((e21 = op2(e21, e21)) & (e21 = op2(e21, e21)) & (e20 = op2(e21, e20)) & (e20 = op2(e20, e21))) | ((e21 = op2(e21, e20)) & (e21 = op2(e20, e21)) & (e20 = op2(e21, e20)) & (e20 = op2(e20, e21))) | ((e20 = op2(e20, e23)) & (e20 = op2(e23, e20)) & (e20 = op2(e20, e20)) & (e20 = op2(e20, e20))) | ((e20 = op2(e20, e22)) & (e20 = op2(e22, e20)) & (e20 = op2(e20, e20)) & (e20 = op2(e20, e20))) | ((e20 = op2(e20, e21)) & (e20 = op2(e21, e20)) & (e20 = op2(e20, e20)) & (e20 = op2(e20, e20))) | ((e20 = op2(e20, e20)) & (e20 = op2(e20, e20)) & (e20 = op2(e20, e20)) & (e20 = op2(e20, e20)))) & ((~ (e23 = op2(e23, op2(e23, e23))) & ~ (e22 = op2(e23, op2(e23, e22))) & ~ (e21 = op2(e23, op2(e23, e21))) & ~ (e20 = op2(e23, op2(e23, e20)))) | (~ (e23 = op2(e22, op2(e22, e23))) & ~ (e22 = op2(e22, op2(e22, e22))) & ~ (e21 = op2(e22, op2(e22, e21))) & ~ (e20 = op2(e22, op2(e22, e20)))) | (~ (e23 = op2(e21, op2(e21, e23))) & ~ (e22 = op2(e21, op2(e21, e22))) & ~ (e21 = op2(e21, op2(e21, e21))) & ~ (e20 = op2(e21, op2(e21, e20)))) | (~ (e23 = op2(e20, op2(e20, e23))) & ~ (e22 = op2(e20, op2(e20, e22))) & ~ (e21 = op2(e20, op2(e20, e21))) & ~ (e20 = op2(e20, op2(e20, e20))))) & ((e23 = op2(e23, e23)) | (e23 = op2(e22, e22)) | (e23 = op2(e21, e21)) | (op2(e20, e20) = e23)) & ((e22 = op2(e23, e23)) | (e22 = op2(e22, e22)) | (e22 = op2(e21, e21)) | (op2(e20, e20) = e22)) & ((e21 = op2(e23, e23)) | (e21 = op2(e22, e22)) | (e21 = op2(e21, e21)) | (op2(e20, e20) = e21)) & ((e20 = op2(e23, e23)) | (e20 = op2(e22, e22)) | (e20 = op2(e21, e21)) | (e20 = op2(e20, e20)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG115+1.p', ax11)).
fof(f2983, plain, (spl144_127 | spl144_107 | spl144_87 | spl144_67), inference(avatar_split_clause, [], [f1074, f1483, f1568, f1653, f1738])).
fof(f1074, plain, ((e22 = op2(e23, e23)) | (e22 = op2(e22, e22)) | (e22 = op2(e21, e21)) | (op2(e20, e20) = e22)), inference(cnf_transformation, [], [f154])).
fof(f2982, plain, (spl144_128 | spl144_108 | spl144_88 | spl144_68), inference(avatar_split_clause, [], [f1075, f1487, f1572, f1657, f1742])).
fof(f1075, plain, ((e23 = op2(e23, e23)) | (e23 = op2(e22, e22)) | (e23 = op2(e21, e21)) | (op2(e20, e20) = e23)), inference(cnf_transformation, [], [f154])).
fof(f2381, plain, (spl144_61 | spl144_41 | spl144_21 | spl144_1), inference(avatar_split_clause, [], [f796, f1171, f1256, f1341, f1426])).
fof(f796, plain, ((e10 = op1(e13, e13)) | (e10 = op1(e12, e12)) | (e10 = op1(e11, e11)) | (e10 = op1(e10, e10))), inference(cnf_transformation, [], [f87])).
fof(f87, plain, ((((e13 = op1(e13, e13)) & (e13 = op1(e13, e13)) & (e13 = op1(e13, e13)) & (e13 = op1(e13, e13))) | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3) & ((~ (e13 = op1(e13, op1(e13, e13))) & ~ (e12 = op1(e13, op1(e13, e12))) & ~ (e11 = op1(e13, op1(e13, e11))) & ~ (e10 = op1(e13, op1(e13, e10)))) | sP2 | sP1 | sP0) & ((e13 = op1(e13, e13)) | (e13 = op1(e12, e12)) | (e13 = op1(e11, e11)) | (op1(e10, e10) = e13)) & ((e12 = op1(e13, e13)) | (e12 = op1(e12, e12)) | (e12 = op1(e11, e11)) | (op1(e10, e10) = e12)) & ((e11 = op1(e13, e13)) | (e11 = op1(e12, e12)) | (e11 = op1(e11, e11)) | (op1(e10, e10) = e11)) & ((e10 = op1(e13, e13)) | (e10 = op1(e12, e12)) | (e10 = op1(e11, e11)) | (e10 = op1(e10, e10)))), inference(definition_folding, [], [f10, e86, e85, e84, e83, e82, e81, e80, e79, e78, e77, e76, e75, e74, e73, e72, e71, e70, e69, e68, e67, e66, e65, e64, e63, e62, e61, e60, e59, e58, e57, e56, e55, e54, e53, e52, e51, e50, e49, e48, e47, e46, e45, e44, e43, e42, e41, e40, e39, e38, e37, e36, e35, e34, e33, e32, e31, e30, e29, e28, e27, e26, e25, e24, e23, e22, e21])).
fof(f21, plain, ((~ (e13 = op1(e10, op1(e10, e13))) & ~ (e12 = op1(e10, op1(e10, e12))) & ~ (e11 = op1(e10, op1(e10, e11))) & ~ (e10 = op1(e10, op1(e10, e10)))) | ~ sP0), inference(usedef, [], [e21])).
fof(e21, plain, (sP0 <=> (~ (e13 = op1(e10, op1(e10, e13))) & ~ (e12 = op1(e10, op1(e10, e12))) & ~ (e11 = op1(e10, op1(e10, e11))) & ~ (e10 = op1(e10, op1(e10, e10))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f22, plain, ((~ (e13 = op1(e11, op1(e11, e13))) & ~ (e12 = op1(e11, op1(e11, e12))) & ~ (e11 = op1(e11, op1(e11, e11))) & ~ (e10 = op1(e11, op1(e11, e10)))) | ~ sP1), inference(usedef, [], [e22])).
fof(e22, plain, (sP1 <=> (~ (e13 = op1(e11, op1(e11, e13))) & ~ (e12 = op1(e11, op1(e11, e12))) & ~ (e11 = op1(e11, op1(e11, e11))) & ~ (e10 = op1(e11, op1(e11, e10))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f23, plain, ((~ (e13 = op1(e12, op1(e12, e13))) & ~ (e12 = op1(e12, op1(e12, e12))) & ~ (e11 = op1(e12, op1(e12, e11))) & ~ (e10 = op1(e12, op1(e12, e10)))) | ~ sP2), inference(usedef, [], [e23])).
fof(e23, plain, (sP2 <=> (~ (e13 = op1(e12, op1(e12, e13))) & ~ (e12 = op1(e12, op1(e12, e12))) & ~ (e11 = op1(e12, op1(e12, e11))) & ~ (e10 = op1(e12, op1(e12, e10))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f24, plain, (((e10 = op1(e10, e10)) & (e10 = op1(e10, e10)) & (e10 = op1(e10, e10)) & (e10 = op1(e10, e10))) | ~ sP3), inference(usedef, [], [e24])).
fof(e24, plain, (sP3 <=> ((e10 = op1(e10, e10)) & (e10 = op1(e10, e10)) & (e10 = op1(e10, e10)) & (e10 = op1(e10, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f25, plain, (((e10 = op1(e10, e11)) & (e10 = op1(e11, e10)) & (e10 = op1(e10, e10)) & (e10 = op1(e10, e10))) | ~ sP4), inference(usedef, [], [e25])).
fof(e25, plain, (sP4 <=> ((e10 = op1(e10, e11)) & (e10 = op1(e11, e10)) & (e10 = op1(e10, e10)) & (e10 = op1(e10, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f26, plain, (((e10 = op1(e10, e12)) & (e10 = op1(e12, e10)) & (e10 = op1(e10, e10)) & (e10 = op1(e10, e10))) | ~ sP5), inference(usedef, [], [e26])).
fof(e26, plain, (sP5 <=> ((e10 = op1(e10, e12)) & (e10 = op1(e12, e10)) & (e10 = op1(e10, e10)) & (e10 = op1(e10, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f27, plain, (((e10 = op1(e10, e13)) & (e10 = op1(e13, e10)) & (e10 = op1(e10, e10)) & (e10 = op1(e10, e10))) | ~ sP6), inference(usedef, [], [e27])).
fof(e27, plain, (sP6 <=> ((e10 = op1(e10, e13)) & (e10 = op1(e13, e10)) & (e10 = op1(e10, e10)) & (e10 = op1(e10, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f28, plain, (((e11 = op1(e11, e10)) & (e11 = op1(e10, e11)) & (e10 = op1(e11, e10)) & (e10 = op1(e10, e11))) | ~ sP7), inference(usedef, [], [e28])).
fof(e28, plain, (sP7 <=> ((e11 = op1(e11, e10)) & (e11 = op1(e10, e11)) & (e10 = op1(e11, e10)) & (e10 = op1(e10, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f29, plain, (((e11 = op1(e11, e11)) & (e11 = op1(e11, e11)) & (e10 = op1(e11, e10)) & (e10 = op1(e10, e11))) | ~ sP8), inference(usedef, [], [e29])).
fof(e29, plain, (sP8 <=> ((e11 = op1(e11, e11)) & (e11 = op1(e11, e11)) & (e10 = op1(e11, e10)) & (e10 = op1(e10, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f30, plain, (((e11 = op1(e11, e12)) & (e11 = op1(e12, e11)) & (e10 = op1(e11, e10)) & (e10 = op1(e10, e11))) | ~ sP9), inference(usedef, [], [e30])).
fof(e30, plain, (sP9 <=> ((e11 = op1(e11, e12)) & (e11 = op1(e12, e11)) & (e10 = op1(e11, e10)) & (e10 = op1(e10, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f31, plain, (((e11 = op1(e11, e13)) & (e11 = op1(e13, e11)) & (e10 = op1(e11, e10)) & (e10 = op1(e10, e11))) | ~ sP10), inference(usedef, [], [e31])).
fof(e31, plain, (sP10 <=> ((e11 = op1(e11, e13)) & (e11 = op1(e13, e11)) & (e10 = op1(e11, e10)) & (e10 = op1(e10, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f32, plain, (((e12 = op1(e12, e10)) & (e12 = op1(e10, e12)) & (e10 = op1(e12, e10)) & (e10 = op1(e10, e12))) | ~ sP11), inference(usedef, [], [e32])).
fof(e32, plain, (sP11 <=> ((e12 = op1(e12, e10)) & (e12 = op1(e10, e12)) & (e10 = op1(e12, e10)) & (e10 = op1(e10, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f33, plain, (((e12 = op1(e12, e11)) & (e12 = op1(e11, e12)) & (e10 = op1(e12, e10)) & (e10 = op1(e10, e12))) | ~ sP12), inference(usedef, [], [e33])).
fof(e33, plain, (sP12 <=> ((e12 = op1(e12, e11)) & (e12 = op1(e11, e12)) & (e10 = op1(e12, e10)) & (e10 = op1(e10, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f34, plain, (((e12 = op1(e12, e12)) & (e12 = op1(e12, e12)) & (e10 = op1(e12, e10)) & (e10 = op1(e10, e12))) | ~ sP13), inference(usedef, [], [e34])).
fof(e34, plain, (sP13 <=> ((e12 = op1(e12, e12)) & (e12 = op1(e12, e12)) & (e10 = op1(e12, e10)) & (e10 = op1(e10, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f35, plain, (((e12 = op1(e12, e13)) & (e12 = op1(e13, e12)) & (e10 = op1(e12, e10)) & (e10 = op1(e10, e12))) | ~ sP14), inference(usedef, [], [e35])).
fof(e35, plain, (sP14 <=> ((e12 = op1(e12, e13)) & (e12 = op1(e13, e12)) & (e10 = op1(e12, e10)) & (e10 = op1(e10, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f36, plain, (((e13 = op1(e13, e10)) & (e13 = op1(e10, e13)) & (e10 = op1(e13, e10)) & (e10 = op1(e10, e13))) | ~ sP15), inference(usedef, [], [e36])).
fof(e36, plain, (sP15 <=> ((e13 = op1(e13, e10)) & (e13 = op1(e10, e13)) & (e10 = op1(e13, e10)) & (e10 = op1(e10, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f37, plain, (((e13 = op1(e13, e11)) & (e13 = op1(e11, e13)) & (e10 = op1(e13, e10)) & (e10 = op1(e10, e13))) | ~ sP16), inference(usedef, [], [e37])).
fof(e37, plain, (sP16 <=> ((e13 = op1(e13, e11)) & (e13 = op1(e11, e13)) & (e10 = op1(e13, e10)) & (e10 = op1(e10, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f38, plain, (((e13 = op1(e13, e12)) & (e13 = op1(e12, e13)) & (e10 = op1(e13, e10)) & (e10 = op1(e10, e13))) | ~ sP17), inference(usedef, [], [e38])).
fof(e38, plain, (sP17 <=> ((e13 = op1(e13, e12)) & (e13 = op1(e12, e13)) & (e10 = op1(e13, e10)) & (e10 = op1(e10, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f39, plain, (((e13 = op1(e13, e13)) & (e13 = op1(e13, e13)) & (e10 = op1(e13, e10)) & (e10 = op1(e10, e13))) | ~ sP18), inference(usedef, [], [e39])).
fof(e39, plain, (sP18 <=> ((e13 = op1(e13, e13)) & (e13 = op1(e13, e13)) & (e10 = op1(e13, e10)) & (e10 = op1(e10, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f40, plain, (((e10 = op1(e10, e10)) & (e10 = op1(e10, e10)) & (e11 = op1(e10, e11)) & (e11 = op1(e11, e10))) | ~ sP19), inference(usedef, [], [e40])).
fof(e40, plain, (sP19 <=> ((e10 = op1(e10, e10)) & (e10 = op1(e10, e10)) & (e11 = op1(e10, e11)) & (e11 = op1(e11, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP19])])).
fof(f41, plain, (((e10 = op1(e10, e11)) & (e10 = op1(e11, e10)) & (e11 = op1(e10, e11)) & (e11 = op1(e11, e10))) | ~ sP20), inference(usedef, [], [e41])).
fof(e41, plain, (sP20 <=> ((e10 = op1(e10, e11)) & (e10 = op1(e11, e10)) & (e11 = op1(e10, e11)) & (e11 = op1(e11, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP20])])).
fof(f42, plain, (((e10 = op1(e10, e12)) & (e10 = op1(e12, e10)) & (e11 = op1(e10, e11)) & (e11 = op1(e11, e10))) | ~ sP21), inference(usedef, [], [e42])).
fof(e42, plain, (sP21 <=> ((e10 = op1(e10, e12)) & (e10 = op1(e12, e10)) & (e11 = op1(e10, e11)) & (e11 = op1(e11, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP21])])).
fof(f43, plain, (((e10 = op1(e10, e13)) & (e10 = op1(e13, e10)) & (e11 = op1(e10, e11)) & (e11 = op1(e11, e10))) | ~ sP22), inference(usedef, [], [e43])).
fof(e43, plain, (sP22 <=> ((e10 = op1(e10, e13)) & (e10 = op1(e13, e10)) & (e11 = op1(e10, e11)) & (e11 = op1(e11, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP22])])).
fof(f44, plain, (((e11 = op1(e11, e10)) & (e11 = op1(e10, e11)) & (e11 = op1(e11, e11)) & (e11 = op1(e11, e11))) | ~ sP23), inference(usedef, [], [e44])).
fof(e44, plain, (sP23 <=> ((e11 = op1(e11, e10)) & (e11 = op1(e10, e11)) & (e11 = op1(e11, e11)) & (e11 = op1(e11, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP23])])).
fof(f45, plain, (((e11 = op1(e11, e11)) & (e11 = op1(e11, e11)) & (e11 = op1(e11, e11)) & (e11 = op1(e11, e11))) | ~ sP24), inference(usedef, [], [e45])).
fof(e45, plain, (sP24 <=> ((e11 = op1(e11, e11)) & (e11 = op1(e11, e11)) & (e11 = op1(e11, e11)) & (e11 = op1(e11, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP24])])).
fof(f46, plain, (((e11 = op1(e11, e12)) & (e11 = op1(e12, e11)) & (e11 = op1(e11, e11)) & (e11 = op1(e11, e11))) | ~ sP25), inference(usedef, [], [e46])).
fof(e46, plain, (sP25 <=> ((e11 = op1(e11, e12)) & (e11 = op1(e12, e11)) & (e11 = op1(e11, e11)) & (e11 = op1(e11, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP25])])).
fof(f47, plain, (((e11 = op1(e11, e13)) & (e11 = op1(e13, e11)) & (e11 = op1(e11, e11)) & (e11 = op1(e11, e11))) | ~ sP26), inference(usedef, [], [e47])).
fof(e47, plain, (sP26 <=> ((e11 = op1(e11, e13)) & (e11 = op1(e13, e11)) & (e11 = op1(e11, e11)) & (e11 = op1(e11, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP26])])).
fof(f48, plain, (((e12 = op1(e12, e10)) & (e12 = op1(e10, e12)) & (e11 = op1(e12, e11)) & (e11 = op1(e11, e12))) | ~ sP27), inference(usedef, [], [e48])).
fof(e48, plain, (sP27 <=> ((e12 = op1(e12, e10)) & (e12 = op1(e10, e12)) & (e11 = op1(e12, e11)) & (e11 = op1(e11, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP27])])).
fof(f49, plain, (((e12 = op1(e12, e11)) & (e12 = op1(e11, e12)) & (e11 = op1(e12, e11)) & (e11 = op1(e11, e12))) | ~ sP28), inference(usedef, [], [e49])).
fof(e49, plain, (sP28 <=> ((e12 = op1(e12, e11)) & (e12 = op1(e11, e12)) & (e11 = op1(e12, e11)) & (e11 = op1(e11, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP28])])).
fof(f50, plain, (((e12 = op1(e12, e12)) & (e12 = op1(e12, e12)) & (e11 = op1(e12, e11)) & (e11 = op1(e11, e12))) | ~ sP29), inference(usedef, [], [e50])).
fof(e50, plain, (sP29 <=> ((e12 = op1(e12, e12)) & (e12 = op1(e12, e12)) & (e11 = op1(e12, e11)) & (e11 = op1(e11, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP29])])).
fof(f51, plain, (((e12 = op1(e12, e13)) & (e12 = op1(e13, e12)) & (e11 = op1(e12, e11)) & (e11 = op1(e11, e12))) | ~ sP30), inference(usedef, [], [e51])).
fof(e51, plain, (sP30 <=> ((e12 = op1(e12, e13)) & (e12 = op1(e13, e12)) & (e11 = op1(e12, e11)) & (e11 = op1(e11, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP30])])).
fof(f52, plain, (((e13 = op1(e13, e10)) & (e13 = op1(e10, e13)) & (e11 = op1(e13, e11)) & (e11 = op1(e11, e13))) | ~ sP31), inference(usedef, [], [e52])).
fof(e52, plain, (sP31 <=> ((e13 = op1(e13, e10)) & (e13 = op1(e10, e13)) & (e11 = op1(e13, e11)) & (e11 = op1(e11, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP31])])).
fof(f53, plain, (((e13 = op1(e13, e11)) & (e13 = op1(e11, e13)) & (e11 = op1(e13, e11)) & (e11 = op1(e11, e13))) | ~ sP32), inference(usedef, [], [e53])).
fof(e53, plain, (sP32 <=> ((e13 = op1(e13, e11)) & (e13 = op1(e11, e13)) & (e11 = op1(e13, e11)) & (e11 = op1(e11, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP32])])).
fof(f54, plain, (((e13 = op1(e13, e12)) & (e13 = op1(e12, e13)) & (e11 = op1(e13, e11)) & (e11 = op1(e11, e13))) | ~ sP33), inference(usedef, [], [e54])).
fof(e54, plain, (sP33 <=> ((e13 = op1(e13, e12)) & (e13 = op1(e12, e13)) & (e11 = op1(e13, e11)) & (e11 = op1(e11, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP33])])).
fof(f55, plain, (((e13 = op1(e13, e13)) & (e13 = op1(e13, e13)) & (e11 = op1(e13, e11)) & (e11 = op1(e11, e13))) | ~ sP34), inference(usedef, [], [e55])).
fof(e55, plain, (sP34 <=> ((e13 = op1(e13, e13)) & (e13 = op1(e13, e13)) & (e11 = op1(e13, e11)) & (e11 = op1(e11, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP34])])).
fof(f56, plain, (((e10 = op1(e10, e10)) & (e10 = op1(e10, e10)) & (e12 = op1(e10, e12)) & (e12 = op1(e12, e10))) | ~ sP35), inference(usedef, [], [e56])).
fof(e56, plain, (sP35 <=> ((e10 = op1(e10, e10)) & (e10 = op1(e10, e10)) & (e12 = op1(e10, e12)) & (e12 = op1(e12, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP35])])).
fof(f57, plain, (((e10 = op1(e10, e11)) & (e10 = op1(e11, e10)) & (e12 = op1(e10, e12)) & (e12 = op1(e12, e10))) | ~ sP36), inference(usedef, [], [e57])).
fof(e57, plain, (sP36 <=> ((e10 = op1(e10, e11)) & (e10 = op1(e11, e10)) & (e12 = op1(e10, e12)) & (e12 = op1(e12, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP36])])).
fof(f58, plain, (((e10 = op1(e10, e12)) & (e10 = op1(e12, e10)) & (e12 = op1(e10, e12)) & (e12 = op1(e12, e10))) | ~ sP37), inference(usedef, [], [e58])).
fof(e58, plain, (sP37 <=> ((e10 = op1(e10, e12)) & (e10 = op1(e12, e10)) & (e12 = op1(e10, e12)) & (e12 = op1(e12, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP37])])).
fof(f59, plain, (((e10 = op1(e10, e13)) & (e10 = op1(e13, e10)) & (e12 = op1(e10, e12)) & (e12 = op1(e12, e10))) | ~ sP38), inference(usedef, [], [e59])).
fof(e59, plain, (sP38 <=> ((e10 = op1(e10, e13)) & (e10 = op1(e13, e10)) & (e12 = op1(e10, e12)) & (e12 = op1(e12, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP38])])).
fof(f60, plain, (((e11 = op1(e11, e10)) & (e11 = op1(e10, e11)) & (e12 = op1(e11, e12)) & (e12 = op1(e12, e11))) | ~ sP39), inference(usedef, [], [e60])).
fof(e60, plain, (sP39 <=> ((e11 = op1(e11, e10)) & (e11 = op1(e10, e11)) & (e12 = op1(e11, e12)) & (e12 = op1(e12, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP39])])).
fof(f61, plain, (((e11 = op1(e11, e11)) & (e11 = op1(e11, e11)) & (e12 = op1(e11, e12)) & (e12 = op1(e12, e11))) | ~ sP40), inference(usedef, [], [e61])).
fof(e61, plain, (sP40 <=> ((e11 = op1(e11, e11)) & (e11 = op1(e11, e11)) & (e12 = op1(e11, e12)) & (e12 = op1(e12, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP40])])).
fof(f62, plain, (((e11 = op1(e11, e12)) & (e11 = op1(e12, e11)) & (e12 = op1(e11, e12)) & (e12 = op1(e12, e11))) | ~ sP41), inference(usedef, [], [e62])).
fof(e62, plain, (sP41 <=> ((e11 = op1(e11, e12)) & (e11 = op1(e12, e11)) & (e12 = op1(e11, e12)) & (e12 = op1(e12, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP41])])).
fof(f63, plain, (((e11 = op1(e11, e13)) & (e11 = op1(e13, e11)) & (e12 = op1(e11, e12)) & (e12 = op1(e12, e11))) | ~ sP42), inference(usedef, [], [e63])).
fof(e63, plain, (sP42 <=> ((e11 = op1(e11, e13)) & (e11 = op1(e13, e11)) & (e12 = op1(e11, e12)) & (e12 = op1(e12, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP42])])).
fof(f64, plain, (((e12 = op1(e12, e10)) & (e12 = op1(e10, e12)) & (e12 = op1(e12, e12)) & (e12 = op1(e12, e12))) | ~ sP43), inference(usedef, [], [e64])).
fof(e64, plain, (sP43 <=> ((e12 = op1(e12, e10)) & (e12 = op1(e10, e12)) & (e12 = op1(e12, e12)) & (e12 = op1(e12, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP43])])).
fof(f65, plain, (((e12 = op1(e12, e11)) & (e12 = op1(e11, e12)) & (e12 = op1(e12, e12)) & (e12 = op1(e12, e12))) | ~ sP44), inference(usedef, [], [e65])).
fof(e65, plain, (sP44 <=> ((e12 = op1(e12, e11)) & (e12 = op1(e11, e12)) & (e12 = op1(e12, e12)) & (e12 = op1(e12, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP44])])).
fof(f66, plain, (((e12 = op1(e12, e12)) & (e12 = op1(e12, e12)) & (e12 = op1(e12, e12)) & (e12 = op1(e12, e12))) | ~ sP45), inference(usedef, [], [e66])).
fof(e66, plain, (sP45 <=> ((e12 = op1(e12, e12)) & (e12 = op1(e12, e12)) & (e12 = op1(e12, e12)) & (e12 = op1(e12, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP45])])).
fof(f67, plain, (((e12 = op1(e12, e13)) & (e12 = op1(e13, e12)) & (e12 = op1(e12, e12)) & (e12 = op1(e12, e12))) | ~ sP46), inference(usedef, [], [e67])).
fof(e67, plain, (sP46 <=> ((e12 = op1(e12, e13)) & (e12 = op1(e13, e12)) & (e12 = op1(e12, e12)) & (e12 = op1(e12, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP46])])).
fof(f68, plain, (((e13 = op1(e13, e10)) & (e13 = op1(e10, e13)) & (e12 = op1(e13, e12)) & (e12 = op1(e12, e13))) | ~ sP47), inference(usedef, [], [e68])).
fof(e68, plain, (sP47 <=> ((e13 = op1(e13, e10)) & (e13 = op1(e10, e13)) & (e12 = op1(e13, e12)) & (e12 = op1(e12, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP47])])).
fof(f69, plain, (((e13 = op1(e13, e11)) & (e13 = op1(e11, e13)) & (e12 = op1(e13, e12)) & (e12 = op1(e12, e13))) | ~ sP48), inference(usedef, [], [e69])).
fof(e69, plain, (sP48 <=> ((e13 = op1(e13, e11)) & (e13 = op1(e11, e13)) & (e12 = op1(e13, e12)) & (e12 = op1(e12, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP48])])).
fof(f70, plain, (((e13 = op1(e13, e12)) & (e13 = op1(e12, e13)) & (e12 = op1(e13, e12)) & (e12 = op1(e12, e13))) | ~ sP49), inference(usedef, [], [e70])).
fof(e70, plain, (sP49 <=> ((e13 = op1(e13, e12)) & (e13 = op1(e12, e13)) & (e12 = op1(e13, e12)) & (e12 = op1(e12, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP49])])).
fof(f71, plain, (((e13 = op1(e13, e13)) & (e13 = op1(e13, e13)) & (e12 = op1(e13, e12)) & (e12 = op1(e12, e13))) | ~ sP50), inference(usedef, [], [e71])).
fof(e71, plain, (sP50 <=> ((e13 = op1(e13, e13)) & (e13 = op1(e13, e13)) & (e12 = op1(e13, e12)) & (e12 = op1(e12, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP50])])).
fof(f72, plain, (((e10 = op1(e10, e10)) & (e10 = op1(e10, e10)) & (e13 = op1(e10, e13)) & (e13 = op1(e13, e10))) | ~ sP51), inference(usedef, [], [e72])).
fof(e72, plain, (sP51 <=> ((e10 = op1(e10, e10)) & (e10 = op1(e10, e10)) & (e13 = op1(e10, e13)) & (e13 = op1(e13, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP51])])).
fof(f73, plain, (((e10 = op1(e10, e11)) & (e10 = op1(e11, e10)) & (e13 = op1(e10, e13)) & (e13 = op1(e13, e10))) | ~ sP52), inference(usedef, [], [e73])).
fof(e73, plain, (sP52 <=> ((e10 = op1(e10, e11)) & (e10 = op1(e11, e10)) & (e13 = op1(e10, e13)) & (e13 = op1(e13, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP52])])).
fof(f74, plain, (((e10 = op1(e10, e12)) & (e10 = op1(e12, e10)) & (e13 = op1(e10, e13)) & (e13 = op1(e13, e10))) | ~ sP53), inference(usedef, [], [e74])).
fof(e74, plain, (sP53 <=> ((e10 = op1(e10, e12)) & (e10 = op1(e12, e10)) & (e13 = op1(e10, e13)) & (e13 = op1(e13, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP53])])).
fof(f75, plain, (((e10 = op1(e10, e13)) & (e10 = op1(e13, e10)) & (e13 = op1(e10, e13)) & (e13 = op1(e13, e10))) | ~ sP54), inference(usedef, [], [e75])).
fof(e75, plain, (sP54 <=> ((e10 = op1(e10, e13)) & (e10 = op1(e13, e10)) & (e13 = op1(e10, e13)) & (e13 = op1(e13, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP54])])).
fof(f76, plain, (((e11 = op1(e11, e10)) & (e11 = op1(e10, e11)) & (e13 = op1(e11, e13)) & (e13 = op1(e13, e11))) | ~ sP55), inference(usedef, [], [e76])).
fof(e76, plain, (sP55 <=> ((e11 = op1(e11, e10)) & (e11 = op1(e10, e11)) & (e13 = op1(e11, e13)) & (e13 = op1(e13, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP55])])).
fof(f77, plain, (((e11 = op1(e11, e11)) & (e11 = op1(e11, e11)) & (e13 = op1(e11, e13)) & (e13 = op1(e13, e11))) | ~ sP56), inference(usedef, [], [e77])).
fof(e77, plain, (sP56 <=> ((e11 = op1(e11, e11)) & (e11 = op1(e11, e11)) & (e13 = op1(e11, e13)) & (e13 = op1(e13, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP56])])).
fof(f78, plain, (((e11 = op1(e11, e12)) & (e11 = op1(e12, e11)) & (e13 = op1(e11, e13)) & (e13 = op1(e13, e11))) | ~ sP57), inference(usedef, [], [e78])).
fof(e78, plain, (sP57 <=> ((e11 = op1(e11, e12)) & (e11 = op1(e12, e11)) & (e13 = op1(e11, e13)) & (e13 = op1(e13, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP57])])).
fof(f79, plain, (((e11 = op1(e11, e13)) & (e11 = op1(e13, e11)) & (e13 = op1(e11, e13)) & (e13 = op1(e13, e11))) | ~ sP58), inference(usedef, [], [e79])).
fof(e79, plain, (sP58 <=> ((e11 = op1(e11, e13)) & (e11 = op1(e13, e11)) & (e13 = op1(e11, e13)) & (e13 = op1(e13, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP58])])).
fof(f80, plain, (((e12 = op1(e12, e10)) & (e12 = op1(e10, e12)) & (e13 = op1(e12, e13)) & (e13 = op1(e13, e12))) | ~ sP59), inference(usedef, [], [e80])).
fof(e80, plain, (sP59 <=> ((e12 = op1(e12, e10)) & (e12 = op1(e10, e12)) & (e13 = op1(e12, e13)) & (e13 = op1(e13, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP59])])).
fof(f81, plain, (((e12 = op1(e12, e11)) & (e12 = op1(e11, e12)) & (e13 = op1(e12, e13)) & (e13 = op1(e13, e12))) | ~ sP60), inference(usedef, [], [e81])).
fof(e81, plain, (sP60 <=> ((e12 = op1(e12, e11)) & (e12 = op1(e11, e12)) & (e13 = op1(e12, e13)) & (e13 = op1(e13, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP60])])).
fof(f82, plain, (((e12 = op1(e12, e12)) & (e12 = op1(e12, e12)) & (e13 = op1(e12, e13)) & (e13 = op1(e13, e12))) | ~ sP61), inference(usedef, [], [e82])).
fof(e82, plain, (sP61 <=> ((e12 = op1(e12, e12)) & (e12 = op1(e12, e12)) & (e13 = op1(e12, e13)) & (e13 = op1(e13, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP61])])).
fof(f83, plain, (((e12 = op1(e12, e13)) & (e12 = op1(e13, e12)) & (e13 = op1(e12, e13)) & (e13 = op1(e13, e12))) | ~ sP62), inference(usedef, [], [e83])).
fof(e83, plain, (sP62 <=> ((e12 = op1(e12, e13)) & (e12 = op1(e13, e12)) & (e13 = op1(e12, e13)) & (e13 = op1(e13, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP62])])).
fof(f84, plain, (((e13 = op1(e13, e10)) & (e13 = op1(e10, e13)) & (e13 = op1(e13, e13)) & (e13 = op1(e13, e13))) | ~ sP63), inference(usedef, [], [e84])).
fof(e84, plain, (sP63 <=> ((e13 = op1(e13, e10)) & (e13 = op1(e10, e13)) & (e13 = op1(e13, e13)) & (e13 = op1(e13, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP63])])).
fof(f85, plain, (((e13 = op1(e13, e11)) & (e13 = op1(e11, e13)) & (e13 = op1(e13, e13)) & (e13 = op1(e13, e13))) | ~ sP64), inference(usedef, [], [e85])).
fof(e85, plain, (sP64 <=> ((e13 = op1(e13, e11)) & (e13 = op1(e11, e13)) & (e13 = op1(e13, e13)) & (e13 = op1(e13, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP64])])).
fof(f86, plain, (((e13 = op1(e13, e12)) & (e13 = op1(e12, e13)) & (e13 = op1(e13, e13)) & (e13 = op1(e13, e13))) | ~ sP65), inference(usedef, [], [e86])).
fof(e86, plain, (sP65 <=> ((e13 = op1(e13, e12)) & (e13 = op1(e12, e13)) & (e13 = op1(e13, e13)) & (e13 = op1(e13, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP65])])).
fof(f10, plain, ((((e13 = op1(e13, e13)) & (e13 = op1(e13, e13)) & (e13 = op1(e13, e13)) & (e13 = op1(e13, e13))) | ((e13 = op1(e13, e12)) & (e13 = op1(e12, e13)) & (e13 = op1(e13, e13)) & (e13 = op1(e13, e13))) | ((e13 = op1(e13, e11)) & (e13 = op1(e11, e13)) & (e13 = op1(e13, e13)) & (e13 = op1(e13, e13))) | ((e13 = op1(e13, e10)) & (e13 = op1(e10, e13)) & (e13 = op1(e13, e13)) & (e13 = op1(e13, e13))) | ((e12 = op1(e12, e13)) & (e12 = op1(e13, e12)) & (e13 = op1(e12, e13)) & (e13 = op1(e13, e12))) | ((e12 = op1(e12, e12)) & (e12 = op1(e12, e12)) & (e13 = op1(e12, e13)) & (e13 = op1(e13, e12))) | ((e12 = op1(e12, e11)) & (e12 = op1(e11, e12)) & (e13 = op1(e12, e13)) & (e13 = op1(e13, e12))) | ((e12 = op1(e12, e10)) & (e12 = op1(e10, e12)) & (e13 = op1(e12, e13)) & (e13 = op1(e13, e12))) | ((e11 = op1(e11, e13)) & (e11 = op1(e13, e11)) & (e13 = op1(e11, e13)) & (e13 = op1(e13, e11))) | ((e11 = op1(e11, e12)) & (e11 = op1(e12, e11)) & (e13 = op1(e11, e13)) & (e13 = op1(e13, e11))) | ((e11 = op1(e11, e11)) & (e11 = op1(e11, e11)) & (e13 = op1(e11, e13)) & (e13 = op1(e13, e11))) | ((e11 = op1(e11, e10)) & (e11 = op1(e10, e11)) & (e13 = op1(e11, e13)) & (e13 = op1(e13, e11))) | ((e10 = op1(e10, e13)) & (e10 = op1(e13, e10)) & (e13 = op1(e10, e13)) & (e13 = op1(e13, e10))) | ((e10 = op1(e10, e12)) & (e10 = op1(e12, e10)) & (e13 = op1(e10, e13)) & (e13 = op1(e13, e10))) | ((e10 = op1(e10, e11)) & (e10 = op1(e11, e10)) & (e13 = op1(e10, e13)) & (e13 = op1(e13, e10))) | ((e10 = op1(e10, e10)) & (e10 = op1(e10, e10)) & (e13 = op1(e10, e13)) & (e13 = op1(e13, e10))) | ((e13 = op1(e13, e13)) & (e13 = op1(e13, e13)) & (e12 = op1(e13, e12)) & (e12 = op1(e12, e13))) | ((e13 = op1(e13, e12)) & (e13 = op1(e12, e13)) & (e12 = op1(e13, e12)) & (e12 = op1(e12, e13))) | ((e13 = op1(e13, e11)) & (e13 = op1(e11, e13)) & (e12 = op1(e13, e12)) & (e12 = op1(e12, e13))) | ((e13 = op1(e13, e10)) & (e13 = op1(e10, e13)) & (e12 = op1(e13, e12)) & (e12 = op1(e12, e13))) | ((e12 = op1(e12, e13)) & (e12 = op1(e13, e12)) & (e12 = op1(e12, e12)) & (e12 = op1(e12, e12))) | ((e12 = op1(e12, e12)) & (e12 = op1(e12, e12)) & (e12 = op1(e12, e12)) & (e12 = op1(e12, e12))) | ((e12 = op1(e12, e11)) & (e12 = op1(e11, e12)) & (e12 = op1(e12, e12)) & (e12 = op1(e12, e12))) | ((e12 = op1(e12, e10)) & (e12 = op1(e10, e12)) & (e12 = op1(e12, e12)) & (e12 = op1(e12, e12))) | ((e11 = op1(e11, e13)) & (e11 = op1(e13, e11)) & (e12 = op1(e11, e12)) & (e12 = op1(e12, e11))) | ((e11 = op1(e11, e12)) & (e11 = op1(e12, e11)) & (e12 = op1(e11, e12)) & (e12 = op1(e12, e11))) | ((e11 = op1(e11, e11)) & (e11 = op1(e11, e11)) & (e12 = op1(e11, e12)) & (e12 = op1(e12, e11))) | ((e11 = op1(e11, e10)) & (e11 = op1(e10, e11)) & (e12 = op1(e11, e12)) & (e12 = op1(e12, e11))) | ((e10 = op1(e10, e13)) & (e10 = op1(e13, e10)) & (e12 = op1(e10, e12)) & (e12 = op1(e12, e10))) | ((e10 = op1(e10, e12)) & (e10 = op1(e12, e10)) & (e12 = op1(e10, e12)) & (e12 = op1(e12, e10))) | ((e10 = op1(e10, e11)) & (e10 = op1(e11, e10)) & (e12 = op1(e10, e12)) & (e12 = op1(e12, e10))) | ((e10 = op1(e10, e10)) & (e10 = op1(e10, e10)) & (e12 = op1(e10, e12)) & (e12 = op1(e12, e10))) | ((e13 = op1(e13, e13)) & (e13 = op1(e13, e13)) & (e11 = op1(e13, e11)) & (e11 = op1(e11, e13))) | ((e13 = op1(e13, e12)) & (e13 = op1(e12, e13)) & (e11 = op1(e13, e11)) & (e11 = op1(e11, e13))) | ((e13 = op1(e13, e11)) & (e13 = op1(e11, e13)) & (e11 = op1(e13, e11)) & (e11 = op1(e11, e13))) | ((e13 = op1(e13, e10)) & (e13 = op1(e10, e13)) & (e11 = op1(e13, e11)) & (e11 = op1(e11, e13))) | ((e12 = op1(e12, e13)) & (e12 = op1(e13, e12)) & (e11 = op1(e12, e11)) & (e11 = op1(e11, e12))) | ((e12 = op1(e12, e12)) & (e12 = op1(e12, e12)) & (e11 = op1(e12, e11)) & (e11 = op1(e11, e12))) | ((e12 = op1(e12, e11)) & (e12 = op1(e11, e12)) & (e11 = op1(e12, e11)) & (e11 = op1(e11, e12))) | ((e12 = op1(e12, e10)) & (e12 = op1(e10, e12)) & (e11 = op1(e12, e11)) & (e11 = op1(e11, e12))) | ((e11 = op1(e11, e13)) & (e11 = op1(e13, e11)) & (e11 = op1(e11, e11)) & (e11 = op1(e11, e11))) | ((e11 = op1(e11, e12)) & (e11 = op1(e12, e11)) & (e11 = op1(e11, e11)) & (e11 = op1(e11, e11))) | ((e11 = op1(e11, e11)) & (e11 = op1(e11, e11)) & (e11 = op1(e11, e11)) & (e11 = op1(e11, e11))) | ((e11 = op1(e11, e10)) & (e11 = op1(e10, e11)) & (e11 = op1(e11, e11)) & (e11 = op1(e11, e11))) | ((e10 = op1(e10, e13)) & (e10 = op1(e13, e10)) & (e11 = op1(e10, e11)) & (e11 = op1(e11, e10))) | ((e10 = op1(e10, e12)) & (e10 = op1(e12, e10)) & (e11 = op1(e10, e11)) & (e11 = op1(e11, e10))) | ((e10 = op1(e10, e11)) & (e10 = op1(e11, e10)) & (e11 = op1(e10, e11)) & (e11 = op1(e11, e10))) | ((e10 = op1(e10, e10)) & (e10 = op1(e10, e10)) & (e11 = op1(e10, e11)) & (e11 = op1(e11, e10))) | ((e13 = op1(e13, e13)) & (e13 = op1(e13, e13)) & (e10 = op1(e13, e10)) & (e10 = op1(e10, e13))) | ((e13 = op1(e13, e12)) & (e13 = op1(e12, e13)) & (e10 = op1(e13, e10)) & (e10 = op1(e10, e13))) | ((e13 = op1(e13, e11)) & (e13 = op1(e11, e13)) & (e10 = op1(e13, e10)) & (e10 = op1(e10, e13))) | ((e13 = op1(e13, e10)) & (e13 = op1(e10, e13)) & (e10 = op1(e13, e10)) & (e10 = op1(e10, e13))) | ((e12 = op1(e12, e13)) & (e12 = op1(e13, e12)) & (e10 = op1(e12, e10)) & (e10 = op1(e10, e12))) | ((e12 = op1(e12, e12)) & (e12 = op1(e12, e12)) & (e10 = op1(e12, e10)) & (e10 = op1(e10, e12))) | ((e12 = op1(e12, e11)) & (e12 = op1(e11, e12)) & (e10 = op1(e12, e10)) & (e10 = op1(e10, e12))) | ((e12 = op1(e12, e10)) & (e12 = op1(e10, e12)) & (e10 = op1(e12, e10)) & (e10 = op1(e10, e12))) | ((e11 = op1(e11, e13)) & (e11 = op1(e13, e11)) & (e10 = op1(e11, e10)) & (e10 = op1(e10, e11))) | ((e11 = op1(e11, e12)) & (e11 = op1(e12, e11)) & (e10 = op1(e11, e10)) & (e10 = op1(e10, e11))) | ((e11 = op1(e11, e11)) & (e11 = op1(e11, e11)) & (e10 = op1(e11, e10)) & (e10 = op1(e10, e11))) | ((e11 = op1(e11, e10)) & (e11 = op1(e10, e11)) & (e10 = op1(e11, e10)) & (e10 = op1(e10, e11))) | ((e10 = op1(e10, e13)) & (e10 = op1(e13, e10)) & (e10 = op1(e10, e10)) & (e10 = op1(e10, e10))) | ((e10 = op1(e10, e12)) & (e10 = op1(e12, e10)) & (e10 = op1(e10, e10)) & (e10 = op1(e10, e10))) | ((e10 = op1(e10, e11)) & (e10 = op1(e11, e10)) & (e10 = op1(e10, e10)) & (e10 = op1(e10, e10))) | ((e10 = op1(e10, e10)) & (e10 = op1(e10, e10)) & (e10 = op1(e10, e10)) & (e10 = op1(e10, e10)))) & ((~ (e13 = op1(e13, op1(e13, e13))) & ~ (e12 = op1(e13, op1(e13, e12))) & ~ (e11 = op1(e13, op1(e13, e11))) & ~ (e10 = op1(e13, op1(e13, e10)))) | (~ (e13 = op1(e12, op1(e12, e13))) & ~ (e12 = op1(e12, op1(e12, e12))) & ~ (e11 = op1(e12, op1(e12, e11))) & ~ (e10 = op1(e12, op1(e12, e10)))) | (~ (e13 = op1(e11, op1(e11, e13))) & ~ (e12 = op1(e11, op1(e11, e12))) & ~ (e11 = op1(e11, op1(e11, e11))) & ~ (e10 = op1(e11, op1(e11, e10)))) | (~ (e13 = op1(e10, op1(e10, e13))) & ~ (e12 = op1(e10, op1(e10, e12))) & ~ (e11 = op1(e10, op1(e10, e11))) & ~ (e10 = op1(e10, op1(e10, e10))))) & ((e13 = op1(e13, e13)) | (e13 = op1(e12, e12)) | (e13 = op1(e11, e11)) | (op1(e10, e10) = e13)) & ((e12 = op1(e13, e13)) | (e12 = op1(e12, e12)) | (e12 = op1(e11, e11)) | (op1(e10, e10) = e12)) & ((e11 = op1(e13, e13)) | (e11 = op1(e12, e12)) | (e11 = op1(e11, e11)) | (op1(e10, e10) = e11)) & ((e10 = op1(e13, e13)) | (e10 = op1(e12, e12)) | (e10 = op1(e11, e11)) | (e10 = op1(e10, e10)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG115+1.p', ax10)).
fof(f2379, plain, (spl144_63 | spl144_43 | spl144_23 | spl144_3), inference(avatar_split_clause, [], [f798, f1179, f1264, f1349, f1434])).
fof(f798, plain, ((e12 = op1(e13, e13)) | (e12 = op1(e12, e12)) | (e12 = op1(e11, e11)) | (op1(e10, e10) = e12)), inference(cnf_transformation, [], [f87])).
fof(f2378, plain, (spl144_64 | spl144_44 | spl144_24 | spl144_4), inference(avatar_split_clause, [], [f799, f1183, f1268, f1353, f1438])).
fof(f799, plain, ((e13 = op1(e13, e13)) | (e13 = op1(e12, e12)) | (e13 = op1(e11, e11)) | (op1(e10, e10) = e13)), inference(cnf_transformation, [], [f87])).
fof(f1772, plain, (spl144_127 | spl144_111 | spl144_95 | spl144_79), inference(avatar_split_clause, [], [f381, f1534, f1602, f1670, f1738])).
fof(f381, plain, ((e22 = op2(e23, e20)) | (e22 = op2(e22, e20)) | (e22 = op2(e21, e20)) | (op2(e20, e20) = e22)), inference(cnf_transformation, [], [f4])).
fof(f4, plain, (((e23 = op2(e23, e23)) | (e23 = op2(e22, e23)) | (e23 = op2(e21, e23)) | (e23 = op2(e20, e23))) & ((e23 = op2(e23, e23)) | (e23 = op2(e23, e22)) | (e23 = op2(e23, e21)) | (e23 = op2(e23, e20))) & ((e22 = op2(e23, e23)) | (e22 = op2(e22, e23)) | (e22 = op2(e21, e23)) | (e22 = op2(e20, e23))) & ((e22 = op2(e23, e23)) | (e22 = op2(e23, e22)) | (e22 = op2(e23, e21)) | (e22 = op2(e23, e20))) & ((e21 = op2(e23, e23)) | (e21 = op2(e22, e23)) | (e21 = op2(e21, e23)) | (e21 = op2(e20, e23))) & ((e21 = op2(e23, e23)) | (e21 = op2(e23, e22)) | (e21 = op2(e23, e21)) | (e21 = op2(e23, e20))) & ((e20 = op2(e23, e23)) | (e20 = op2(e22, e23)) | (e20 = op2(e21, e23)) | (e20 = op2(e20, e23))) & ((e20 = op2(e23, e23)) | (e20 = op2(e23, e22)) | (e20 = op2(e23, e21)) | (e20 = op2(e23, e20))) & ((e23 = op2(e23, e22)) | (e23 = op2(e22, e22)) | (e23 = op2(e21, e22)) | (e23 = op2(e20, e22))) & ((e23 = op2(e22, e23)) | (e23 = op2(e22, e22)) | (e23 = op2(e22, e21)) | (e23 = op2(e22, e20))) & ((e22 = op2(e23, e22)) | (e22 = op2(e22, e22)) | (e22 = op2(e21, e22)) | (e22 = op2(e20, e22))) & ((e22 = op2(e22, e23)) | (e22 = op2(e22, e22)) | (e22 = op2(e22, e21)) | (e22 = op2(e22, e20))) & ((e21 = op2(e23, e22)) | (e21 = op2(e22, e22)) | (e21 = op2(e21, e22)) | (e21 = op2(e20, e22))) & ((e21 = op2(e22, e23)) | (e21 = op2(e22, e22)) | (e21 = op2(e22, e21)) | (e21 = op2(e22, e20))) & ((e20 = op2(e23, e22)) | (e20 = op2(e22, e22)) | (e20 = op2(e21, e22)) | (e20 = op2(e20, e22))) & ((e20 = op2(e22, e23)) | (e20 = op2(e22, e22)) | (e20 = op2(e22, e21)) | (e20 = op2(e22, e20))) & ((e23 = op2(e23, e21)) | (e23 = op2(e22, e21)) | (e23 = op2(e21, e21)) | (e23 = op2(e20, e21))) & ((e23 = op2(e21, e23)) | (e23 = op2(e21, e22)) | (e23 = op2(e21, e21)) | (e23 = op2(e21, e20))) & ((e22 = op2(e23, e21)) | (e22 = op2(e22, e21)) | (e22 = op2(e21, e21)) | (e22 = op2(e20, e21))) & ((e22 = op2(e21, e23)) | (e22 = op2(e21, e22)) | (e22 = op2(e21, e21)) | (e22 = op2(e21, e20))) & ((e21 = op2(e23, e21)) | (e21 = op2(e22, e21)) | (e21 = op2(e21, e21)) | (e21 = op2(e20, e21))) & ((e21 = op2(e21, e23)) | (e21 = op2(e21, e22)) | (e21 = op2(e21, e21)) | (e21 = op2(e21, e20))) & ((e20 = op2(e23, e21)) | (e20 = op2(e22, e21)) | (e20 = op2(e21, e21)) | (e20 = op2(e20, e21))) & ((e20 = op2(e21, e23)) | (e20 = op2(e21, e22)) | (e20 = op2(e21, e21)) | (e20 = op2(e21, e20))) & ((e23 = op2(e23, e20)) | (e23 = op2(e22, e20)) | (e23 = op2(e21, e20)) | (op2(e20, e20) = e23)) & ((e23 = op2(e20, e23)) | (e23 = op2(e20, e22)) | (e23 = op2(e20, e21)) | (op2(e20, e20) = e23)) & ((e22 = op2(e23, e20)) | (e22 = op2(e22, e20)) | (e22 = op2(e21, e20)) | (op2(e20, e20) = e22)) & ((e22 = op2(e20, e23)) | (e22 = op2(e20, e22)) | (e22 = op2(e20, e21)) | (op2(e20, e20) = e22)) & ((e21 = op2(e23, e20)) | (e21 = op2(e22, e20)) | (e21 = op2(e21, e20)) | (op2(e20, e20) = e21)) & ((e21 = op2(e20, e23)) | (e21 = op2(e20, e22)) | (e21 = op2(e20, e21)) | (op2(e20, e20) = e21)) & ((e20 = op2(e23, e20)) | (e20 = op2(e22, e20)) | (e20 = op2(e21, e20)) | (e20 = op2(e20, e20))) & ((e20 = op2(e20, e23)) | (e20 = op2(e20, e22)) | (e20 = op2(e20, e21)) | (e20 = op2(e20, e20)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG115+1.p', ax4)).
fof(f1771, plain, (spl144_128 | spl144_124 | spl144_120 | spl144_116), inference(avatar_split_clause, [], [f382, f1691, f1708, f1725, f1742])).
fof(f382, plain, ((e23 = op2(e20, e23)) | (e23 = op2(e20, e22)) | (e23 = op2(e20, e21)) | (op2(e20, e20) = e23)), inference(cnf_transformation, [], [f4])).
fof(f1764, plain, (spl144_123 | spl144_107 | spl144_91 | spl144_75), inference(avatar_split_clause, [], [f389, f1517, f1585, f1653, f1721])).
fof(f389, plain, ((e22 = op2(e23, e21)) | (e22 = op2(e22, e21)) | (e22 = op2(e21, e21)) | (e22 = op2(e20, e21))), inference(cnf_transformation, [], [f4])).
fof(f1761, plain, (spl144_93 | spl144_89 | spl144_85 | spl144_81), inference(avatar_split_clause, [], [f392, f1543, f1560, f1577, f1594])).
fof(f392, plain, ((e20 = op2(e22, e23)) | (e20 = op2(e22, e22)) | (e20 = op2(e22, e21)) | (e20 = op2(e22, e20))), inference(cnf_transformation, [], [f4])).
fof(f1759, plain, (spl144_94 | spl144_90 | spl144_86 | spl144_82), inference(avatar_split_clause, [], [f394, f1547, f1564, f1581, f1598])).
fof(f394, plain, ((e21 = op2(e22, e23)) | (e21 = op2(e22, e22)) | (e21 = op2(e22, e21)) | (e21 = op2(e22, e20))), inference(cnf_transformation, [], [f4])).
fof(f1758, plain, (spl144_118 | spl144_102 | spl144_86 | spl144_70), inference(avatar_split_clause, [], [f395, f1496, f1564, f1632, f1700])).
fof(f395, plain, ((e21 = op2(e23, e22)) | (e21 = op2(e22, e22)) | (e21 = op2(e21, e22)) | (e21 = op2(e20, e22))), inference(cnf_transformation, [], [f4])).
fof(f1755, plain, (spl144_96 | spl144_92 | spl144_88 | spl144_84), inference(avatar_split_clause, [], [f398, f1555, f1572, f1589, f1606])).
fof(f398, plain, ((e23 = op2(e22, e23)) | (e23 = op2(e22, e22)) | (e23 = op2(e22, e21)) | (e23 = op2(e22, e20))), inference(cnf_transformation, [], [f4])).
fof(f1754, plain, (spl144_120 | spl144_104 | spl144_88 | spl144_72), inference(avatar_split_clause, [], [f399, f1504, f1572, f1640, f1708])).
fof(f399, plain, ((e23 = op2(e23, e22)) | (e23 = op2(e22, e22)) | (e23 = op2(e21, e22)) | (e23 = op2(e20, e22))), inference(cnf_transformation, [], [f4])).
fof(f1753, plain, (spl144_77 | spl144_73 | spl144_69 | spl144_65), inference(avatar_split_clause, [], [f400, f1475, f1492, f1509, f1526])).
fof(f400, plain, ((e20 = op2(e23, e23)) | (e20 = op2(e23, e22)) | (e20 = op2(e23, e21)) | (e20 = op2(e23, e20))), inference(cnf_transformation, [], [f4])).
fof(f1677, plain, (spl144_109 | spl144_110 | spl144_111 | spl144_112), inference(avatar_split_clause, [], [f364, f1674, f1670, f1666, f1662])).
fof(f364, plain, ((e23 = op2(e21, e20)) | (e22 = op2(e21, e20)) | (e21 = op2(e21, e20)) | (e20 = op2(e21, e20))), inference(cnf_transformation, [], [f3])).
fof(f3, plain, (((e23 = op2(e23, e23)) | (e22 = op2(e23, e23)) | (e21 = op2(e23, e23)) | (e20 = op2(e23, e23))) & ((e23 = op2(e23, e22)) | (e22 = op2(e23, e22)) | (e21 = op2(e23, e22)) | (e20 = op2(e23, e22))) & ((e23 = op2(e23, e21)) | (e22 = op2(e23, e21)) | (e21 = op2(e23, e21)) | (e20 = op2(e23, e21))) & ((e23 = op2(e23, e20)) | (e22 = op2(e23, e20)) | (e21 = op2(e23, e20)) | (e20 = op2(e23, e20))) & ((e23 = op2(e22, e23)) | (e22 = op2(e22, e23)) | (e21 = op2(e22, e23)) | (e20 = op2(e22, e23))) & ((e23 = op2(e22, e22)) | (e22 = op2(e22, e22)) | (e21 = op2(e22, e22)) | (e20 = op2(e22, e22))) & ((e23 = op2(e22, e21)) | (e22 = op2(e22, e21)) | (e21 = op2(e22, e21)) | (e20 = op2(e22, e21))) & ((e23 = op2(e22, e20)) | (e22 = op2(e22, e20)) | (e21 = op2(e22, e20)) | (e20 = op2(e22, e20))) & ((e23 = op2(e21, e23)) | (e22 = op2(e21, e23)) | (e21 = op2(e21, e23)) | (e20 = op2(e21, e23))) & ((e23 = op2(e21, e22)) | (e22 = op2(e21, e22)) | (e21 = op2(e21, e22)) | (e20 = op2(e21, e22))) & ((e23 = op2(e21, e21)) | (e22 = op2(e21, e21)) | (e21 = op2(e21, e21)) | (e20 = op2(e21, e21))) & ((e23 = op2(e21, e20)) | (e22 = op2(e21, e20)) | (e21 = op2(e21, e20)) | (e20 = op2(e21, e20))) & ((e23 = op2(e20, e23)) | (e22 = op2(e20, e23)) | (e21 = op2(e20, e23)) | (e20 = op2(e20, e23))) & ((e23 = op2(e20, e22)) | (e22 = op2(e20, e22)) | (e21 = op2(e20, e22)) | (e20 = op2(e20, e22))) & ((e23 = op2(e20, e21)) | (e22 = op2(e20, e21)) | (e21 = op2(e20, e21)) | (e20 = op2(e20, e21))) & ((op2(e20, e20) = e23) | (op2(e20, e20) = e22) | (op2(e20, e20) = e21) | (e20 = op2(e20, e20)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG115+1.p', ax3)).
fof(f1469, plain, (spl144_63 | spl144_59 | spl144_55 | spl144_51), inference(avatar_split_clause, [], [f332, f1383, f1400, f1417, f1434])).
fof(f332, plain, ((e12 = op1(e10, e13)) | (e12 = op1(e10, e12)) | (e12 = op1(e10, e11)) | (op1(e10, e10) = e12)), inference(cnf_transformation, [], [f2])).
fof(f2, plain, (((e13 = op1(e13, e13)) | (e13 = op1(e12, e13)) | (e13 = op1(e11, e13)) | (e13 = op1(e10, e13))) & ((e13 = op1(e13, e13)) | (e13 = op1(e13, e12)) | (e13 = op1(e13, e11)) | (e13 = op1(e13, e10))) & ((e12 = op1(e13, e13)) | (e12 = op1(e12, e13)) | (e12 = op1(e11, e13)) | (e12 = op1(e10, e13))) & ((e12 = op1(e13, e13)) | (e12 = op1(e13, e12)) | (e12 = op1(e13, e11)) | (e12 = op1(e13, e10))) & ((e11 = op1(e13, e13)) | (e11 = op1(e12, e13)) | (e11 = op1(e11, e13)) | (e11 = op1(e10, e13))) & ((e11 = op1(e13, e13)) | (e11 = op1(e13, e12)) | (e11 = op1(e13, e11)) | (e11 = op1(e13, e10))) & ((e10 = op1(e13, e13)) | (e10 = op1(e12, e13)) | (e10 = op1(e11, e13)) | (e10 = op1(e10, e13))) & ((e10 = op1(e13, e13)) | (e10 = op1(e13, e12)) | (e10 = op1(e13, e11)) | (e10 = op1(e13, e10))) & ((e13 = op1(e13, e12)) | (e13 = op1(e12, e12)) | (e13 = op1(e11, e12)) | (e13 = op1(e10, e12))) & ((e13 = op1(e12, e13)) | (e13 = op1(e12, e12)) | (e13 = op1(e12, e11)) | (e13 = op1(e12, e10))) & ((e12 = op1(e13, e12)) | (e12 = op1(e12, e12)) | (e12 = op1(e11, e12)) | (e12 = op1(e10, e12))) & ((e12 = op1(e12, e13)) | (e12 = op1(e12, e12)) | (e12 = op1(e12, e11)) | (e12 = op1(e12, e10))) & ((e11 = op1(e13, e12)) | (e11 = op1(e12, e12)) | (e11 = op1(e11, e12)) | (e11 = op1(e10, e12))) & ((e11 = op1(e12, e13)) | (e11 = op1(e12, e12)) | (e11 = op1(e12, e11)) | (e11 = op1(e12, e10))) & ((e10 = op1(e13, e12)) | (e10 = op1(e12, e12)) | (e10 = op1(e11, e12)) | (e10 = op1(e10, e12))) & ((e10 = op1(e12, e13)) | (e10 = op1(e12, e12)) | (e10 = op1(e12, e11)) | (e10 = op1(e12, e10))) & ((e13 = op1(e13, e11)) | (e13 = op1(e12, e11)) | (e13 = op1(e11, e11)) | (e13 = op1(e10, e11))) & ((e13 = op1(e11, e13)) | (e13 = op1(e11, e12)) | (e13 = op1(e11, e11)) | (e13 = op1(e11, e10))) & ((e12 = op1(e13, e11)) | (e12 = op1(e12, e11)) | (e12 = op1(e11, e11)) | (e12 = op1(e10, e11))) & ((e12 = op1(e11, e13)) | (e12 = op1(e11, e12)) | (e12 = op1(e11, e11)) | (e12 = op1(e11, e10))) & ((e11 = op1(e13, e11)) | (e11 = op1(e12, e11)) | (e11 = op1(e11, e11)) | (e11 = op1(e10, e11))) & ((e11 = op1(e11, e13)) | (e11 = op1(e11, e12)) | (e11 = op1(e11, e11)) | (e11 = op1(e11, e10))) & ((e10 = op1(e13, e11)) | (e10 = op1(e12, e11)) | (e10 = op1(e11, e11)) | (e10 = op1(e10, e11))) & ((e10 = op1(e11, e13)) | (e10 = op1(e11, e12)) | (e10 = op1(e11, e11)) | (e10 = op1(e11, e10))) & ((e13 = op1(e13, e10)) | (e13 = op1(e12, e10)) | (e13 = op1(e11, e10)) | (op1(e10, e10) = e13)) & ((e13 = op1(e10, e13)) | (e13 = op1(e10, e12)) | (e13 = op1(e10, e11)) | (op1(e10, e10) = e13)) & ((e12 = op1(e13, e10)) | (e12 = op1(e12, e10)) | (e12 = op1(e11, e10)) | (op1(e10, e10) = e12)) & ((e12 = op1(e10, e13)) | (e12 = op1(e10, e12)) | (e12 = op1(e10, e11)) | (op1(e10, e10) = e12)) & ((e11 = op1(e13, e10)) | (e11 = op1(e12, e10)) | (e11 = op1(e11, e10)) | (op1(e10, e10) = e11)) & ((e11 = op1(e10, e13)) | (e11 = op1(e10, e12)) | (e11 = op1(e10, e11)) | (op1(e10, e10) = e11)) & ((e10 = op1(e13, e10)) | (e10 = op1(e12, e10)) | (e10 = op1(e11, e10)) | (e10 = op1(e10, e10))) & ((e10 = op1(e10, e13)) | (e10 = op1(e10, e12)) | (e10 = op1(e10, e11)) | (e10 = op1(e10, e10)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG115+1.p', ax2)).
fof(f1468, plain, (spl144_63 | spl144_47 | spl144_31 | spl144_15), inference(avatar_split_clause, [], [f333, f1230, f1298, f1366, f1434])).
fof(f333, plain, ((e12 = op1(e13, e10)) | (e12 = op1(e12, e10)) | (e12 = op1(e11, e10)) | (op1(e10, e10) = e12)), inference(cnf_transformation, [], [f2])).
fof(f1466, plain, (spl144_64 | spl144_48 | spl144_32 | spl144_16), inference(avatar_split_clause, [], [f335, f1234, f1302, f1370, f1438])).
fof(f335, plain, ((e13 = op1(e13, e10)) | (e13 = op1(e12, e10)) | (e13 = op1(e11, e10)) | (op1(e10, e10) = e13)), inference(cnf_transformation, [], [f2])).
fof(f1462, plain, (spl144_58 | spl144_42 | spl144_26 | spl144_10), inference(avatar_split_clause, [], [f339, f1209, f1277, f1345, f1413])).
fof(f339, plain, ((e11 = op1(e13, e11)) | (e11 = op1(e12, e11)) | (e11 = op1(e11, e11)) | (e11 = op1(e10, e11))), inference(cnf_transformation, [], [f2])).
fof(f1457, plain, (spl144_29 | spl144_25 | spl144_21 | spl144_17), inference(avatar_split_clause, [], [f344, f1239, f1256, f1273, f1290])).
fof(f344, plain, ((e10 = op1(e12, e13)) | (e10 = op1(e12, e12)) | (e10 = op1(e12, e11)) | (e10 = op1(e12, e10))), inference(cnf_transformation, [], [f2])).
fof(f1454, plain, (spl144_54 | spl144_38 | spl144_22 | spl144_6), inference(avatar_split_clause, [], [f347, f1192, f1260, f1328, f1396])).
fof(f347, plain, ((e11 = op1(e13, e12)) | (e11 = op1(e12, e12)) | (e11 = op1(e11, e12)) | (e11 = op1(e10, e12))), inference(cnf_transformation, [], [f2])).
fof(f1449, plain, (spl144_13 | spl144_9 | spl144_5 | spl144_1), inference(avatar_split_clause, [], [f352, f1171, f1188, f1205, f1222])).
fof(f352, plain, ((e10 = op1(e13, e13)) | (e10 = op1(e13, e12)) | (e10 = op1(e13, e11)) | (e10 = op1(e13, e10))), inference(cnf_transformation, [], [f2])).
fof(f1442, plain, (spl144_52 | spl144_36 | spl144_20 | spl144_4), inference(avatar_split_clause, [], [f359, f1183, f1251, f1319, f1387])).
fof(f359, plain, ((e13 = op1(e13, e13)) | (e13 = op1(e12, e13)) | (e13 = op1(e11, e13)) | (e13 = op1(e10, e13))), inference(cnf_transformation, [], [f2])).
fof(f1373, plain, (spl144_45 | spl144_46 | spl144_47 | spl144_48), inference(avatar_split_clause, [], [f316, f1370, f1366, f1362, f1358])).
fof(f316, plain, ((e13 = op1(e11, e10)) | (e12 = op1(e11, e10)) | (e11 = op1(e11, e10)) | (e10 = op1(e11, e10))), inference(cnf_transformation, [], [f1])).
fof(f1, plain, (((e13 = op1(e13, e13)) | (e12 = op1(e13, e13)) | (e11 = op1(e13, e13)) | (e10 = op1(e13, e13))) & ((e13 = op1(e13, e12)) | (e12 = op1(e13, e12)) | (e11 = op1(e13, e12)) | (e10 = op1(e13, e12))) & ((e13 = op1(e13, e11)) | (e12 = op1(e13, e11)) | (e11 = op1(e13, e11)) | (e10 = op1(e13, e11))) & ((e13 = op1(e13, e10)) | (e12 = op1(e13, e10)) | (e11 = op1(e13, e10)) | (e10 = op1(e13, e10))) & ((e13 = op1(e12, e13)) | (e12 = op1(e12, e13)) | (e11 = op1(e12, e13)) | (e10 = op1(e12, e13))) & ((e13 = op1(e12, e12)) | (e12 = op1(e12, e12)) | (e11 = op1(e12, e12)) | (e10 = op1(e12, e12))) & ((e13 = op1(e12, e11)) | (e12 = op1(e12, e11)) | (e11 = op1(e12, e11)) | (e10 = op1(e12, e11))) & ((e13 = op1(e12, e10)) | (e12 = op1(e12, e10)) | (e11 = op1(e12, e10)) | (e10 = op1(e12, e10))) & ((e13 = op1(e11, e13)) | (e12 = op1(e11, e13)) | (e11 = op1(e11, e13)) | (e10 = op1(e11, e13))) & ((e13 = op1(e11, e12)) | (e12 = op1(e11, e12)) | (e11 = op1(e11, e12)) | (e10 = op1(e11, e12))) & ((e13 = op1(e11, e11)) | (e12 = op1(e11, e11)) | (e11 = op1(e11, e11)) | (e10 = op1(e11, e11))) & ((e13 = op1(e11, e10)) | (e12 = op1(e11, e10)) | (e11 = op1(e11, e10)) | (e10 = op1(e11, e10))) & ((e13 = op1(e10, e13)) | (e12 = op1(e10, e13)) | (e11 = op1(e10, e13)) | (e10 = op1(e10, e13))) & ((e13 = op1(e10, e12)) | (e12 = op1(e10, e12)) | (e11 = op1(e10, e12)) | (e10 = op1(e10, e12))) & ((e13 = op1(e10, e11)) | (e12 = op1(e10, e11)) | (e11 = op1(e10, e11)) | (e10 = op1(e10, e11))) & ((op1(e10, e10) = e13) | (op1(e10, e10) = e12) | (op1(e10, e10) = e11) | (e10 = op1(e10, e10)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG115+1.p', ax1)).
fof(f1203, plain, (spl144_5 | spl144_6 | spl144_7 | spl144_8), inference(avatar_split_clause, [], [f326, f1200, f1196, f1192, f1188])).
fof(f326, plain, ((e13 = op1(e13, e12)) | (e12 = op1(e13, e12)) | (e11 = op1(e13, e12)) | (e10 = op1(e13, e12))), inference(cnf_transformation, [], [f1])).