fof(f19407, plain, $false, inference(avatar_sat_refutation, [], [f712, f733, f754, f775, f796, f817, f838, f859, f880, f901, f922, f943, f964, f985, f1006, f1027, f1048, f1069, f1090, f1111, f1132, f1153, f1174, f1195, f1216, f1237, f1240, f1241, f1242, f1243, f1244, f1245, f1246, f1247, f1248, f1249, f1252, f1253, f1254, f1255, f1256, f1257, f1258, f1259, f1260, f1261, f1264, f1265, f1266, f1267, f1268, f1269, f1270, f1271, f1272, f1273, f1276, f1277, f1278, f1279, f1280, f1281, f1282, f1283, f1284, f1285, f1296, f1301, f1306, f1315, f1320, f1325, f1334, f1339, f1344, f1353, f1358, f1363, f1372, f1377, f1378, f1389, f1398, f1403, f1408, f1417, f1422, f1427, f1436, f1441, f1446, f1455, f1460, f1461, f1470, f1475, f1476, f1487, f1496, f1501, f1506, f1515, f1520, f1525, f1534, f1539, f1540, f1549, f1554, f1555, f1564, f1569, f1570, f1581, f1590, f1595, f1600, f1609, f1614, f1615, f1624, f1629, f1630, f1639, f1644, f1645, f1654, f1659, f1660, f1671, f1678, f1691, f1717, f1726, f1735, f1744, f1749, f1767, f1772, f1781, f1786, f1797, f1798, f1810, f1811, f1813, f1819, f1837, f1874, f1885, f1891, f1896, f1897, f1898, f1899, f1901, f1907, f1912, f1917, f1920, f1922, f1930, f1941, f1947, f1952, f1955, f1956, f1957, f1963, f1975, f1977, f1985, f1986, f1990, f1992, f1993, f1994, f1995, f1996, f1997, f2008, f2013, f2015, f2023, f2026, f2028, f2030, f2031, f2035, f2036, f2037, f2038, f2073, f2077, f2087, f2090, f2101, f2110, f2111, f2121, f2125, f2128, f2131, f2139, f2140, f2141, f2143, f2144, f2145, f2160, f2183, f2189, f2190, f2192, f2201, f2204, f2214, f2215, f2216, f2218, f2225, f2263, f2269, f2277, f2292, f2314, f2320, f2323, f2326, f2333, f2337, f2340, f2343, f2352, f2355, f2362, f2371, f2375, f2379, f2382, f2385, f2388, f2402, f2404, f2411, f2416, f2430, f2436, f2448, f2452, f2457, f2468, f2470, f2472, f2487, f2494, f2500, f2509, f2531, f2532, f2545, f2555, f2556, f2577, f2588, f2595, f2603, f2605, f2612, f2623, f2652, f2658, f2671, f2686, f2711, f2764, f2772, f2775, f2786, f2791, f2808, f2818, f2878, f2881, f2899, f2902, f2905, f2913, f2958, f2961, f3019, f3040, f3045, f3052, f3061, f3062, f3133, f3136, f3179, f3190, f3201, f3212, f3221, f3282, f3301, f3330, f3397, f3409, f3432, f3435, f3439, f3446, f3475, f3487, f3508, f3534, f3544, f3558, f3613, f3618, f3669, f3672, f3697, f3708, f3758, f3835, f3856, f3926, f3931, f4001, f4025, f4035, f4054, f4102, f4105, f4133, f4137, f4150, f4199, f4295, f4330, f4333, f4336, f4342, f4347, f4366, f4418, f4546, f4575, f4618, f4629, f4650, f4652, f4663, f4665, f4742, f4755, f4772, f4787, f4799, f4825, f4913, f4916, f4922, f4980, f4988, f5016, f5025, f5068, f5074, f5084, f5106, f5110, f5112, f5114, f5150, f5191, f5197, f5213, f5216, f5224, f5241, f5244, f5265, f5304, f5369, f5418, f5421, f5433, f5455, f5480, f5526, f5536, f5568, f5573, f5574, f5612, f5617, f5638, f5659, f5665, f5762, f5769, f5780, f5784, f5798, f5836, f5854, f5915, f5920, f5926, f5928, f5942, f5946, f5959, f5974, f5987, f5996, f6026, f6085, f6144, f6161, f6162, f6194, f6196, f6208, f6222, f6232, f6240, f6259, f6274, f6277, f6280, f6319, f6346, f6378, f6398, f6445, f6446, f6461, f6462, f6483, f6537, f6648, f6656, f6679, f6757, f6772, f6777, f6814, f6883, f6898, f6978, f6990, f7027, f7036, f7052, f7058, f7088, f7117, f7164, f7201, f7209, f7338, f7341, f7357, f7366, f7399, f7400, f7420, f7474, f7561, f7563, f7579, f7602, f7638, f7650, f7689, f7721, f7750, f7755, f7772, f7773, f7798, f7801, f7831, f7876, f7979, f7985, f8036, f8057, f8058, f8059, f8070, f8072, f8086, f8089, f8103, f8105, f8116, f8147, f8152, f8162, f8182, f8183, f8185, f8187, f8237, f8343, f8380, f8482, f8529, f8537, f8589, f8599, f8601, f8624, f8638, f8657, f8680, f8686, f8714, f8725, f8767, f8798, f8809, f8810, f8823, f8854, f8859, f8865, f8891, f8902, f8906, f8929, f8966, f9019, f9035, f9037, f9048, f9085, f9094, f9095, f9112, f9134, f9157, f9177, f9179, f9190, f9196, f9232, f9237, f9244, f9291, f9316, f9317, f9318, f9319, f9381, f9385, f9388, f9399, f9403, f9421, f9425, f9432, f9434, f9461, f9493, f9494, f9502, f9504, f9512, f9618, f9624, f9640, f9649, f9659, f9695, f9715, f9775, f9784, f9787, f9795, f9813, f9825, f9868, f9900, f9940, f9955, f9971, f10007, f10012, f10029, f10079, f10085, f10118, f10164, f10204, f10210, f10229, f10233, f10237, f10245, f10253, f10308, f10320, f10321, f10362, f10454, f10474, f10556, f10559, f10572, f10575, f10585, f10593, f10596, f10599, f10601, f10609, f10628, f10632, f10648, f10668, f10669, f10690, f10692, f10739, f10761, f10770, f10772, f10776, f10793, f10805, f10809, f10818, f10827, f10840, f10846, f10861, f10884, f10888, f10890, f10895, f10898, f10938, f10946, f10956, f10991, f11023, f11048, f11079, f11154, f11165, f11167, f11185, f11188, f11238, f11251, f11285, f11295, f11338, f11339, f11342, f11369, f11390, f11403, f11411, f11429, f11430, f11432, f11442, f11489, f11508, f11509, f11511, f11615, f11640, f11652, f11706, f11719, f11735, f11752, f11760, f11770, f11820, f11881, f11890, f11934, f11986, f12004, f12013, f12026, f12036, f12074, f12090, f12097, f12110, f12113, f12136, f12222, f12235, f12245, f12281, f12283, f12285, f12288, f12292, f12333, f12350, f12354, f12361, f12386, f12399, f12474, f12490, f12538, f12600, f12625, f12628, f12653, f12657, f12659, f12677, f12686, f12709, f12710, f12743, f12804, f12812, f12831, f12876, f12877, f12878, f12879, f12881, f12912, f12923, f12992, f12993, f13056, f13093, f13098, f13103, f13105, f13146, f13167, f13168, f13202, f13204, f13207, f13215, f13223, f13227, f13260, f13296, f13329, f13344, f13374, f13385, f13427, f13429, f13442, f13511, f13512, f13533, f13534, f13536, f13558, f13574, f13586, f13587, f13604, f13657, f13698, f13712, f13847, f13849, f13858, f13891, f13904, f13908, f13913, f13979, f13980, f13991, f14018, f14021, f14042, f14049, f14051, f14055, f14062, f14069, f14078, f14080, f14092, f14101, f14104, f14107, f14166, f14171, f14205, f14218, f14229, f14235, f14258, f14328, f14363, f14454, f14477, f14481, f14491, f14492, f14501, f14502, f14503, f14516, f14517, f14532, f14561, f14562, f14589, f14591, f14592, f14602, f14644, f14717, f14724, f14726, f14743, f14773, f14783, f14784, f14826, f14827, f14837, f14844, f14873, f14900, f14927, f14971, f14973, f15060, f15156, f15160, f15194, f15205, f15276, f15283, f15285, f15317, f15318, f15319, f15320, f15321, f15322, f15323, f15324, f15326, f15339, f15350, f15351, f15363, f15376, f15407, f15420, f15461, f15462, f15475, f15513, f15579, f15581, f15593, f15599, f15601, f15656, f15657, f15669, f15670, f15671, f15690, f15721, f15730, f15731, f15746, f15779, f15790, f15800, f15823, f15839, f15841, f15882, f15937, f15944, f15952, f15954, f15989, f15996, f15997, f16014, f16027, f16034, f16085, f16087, f16089, f16107, f16116, f16118, f16126, f16142, f16147, f16152, f16163, f16164, f16166, f16179, f16180, f16193, f16194, f16205, f16211, f16220, f16227, f16240, f16242, f16248, f16256, f16268, f16270, f16337, f16338, f16359, f16361, f16376, f16378, f16386, f16387, f16388, f16389, f16405, f16406, f16418, f16419, f16432, f16454, f16474, f16485, f16504, f16514, f16534, f16543, f16630, f16633, f16637, f16646, f16648, f16655, f16657, f16664, f16669, f16676, f16681, f16683, f16691, f16700, f16708, f16712, f16721, f16727, f16787, f16797, f16798, f16812, f16826, f16827, f16842, f16843, f16852, f16867, f16883, f16887, f16890, f16935, f16941, f16953, f17010, f17011, f17012, f17013, f17014, f17015, f17016, f17017, f17018, f17042, f17043, f17080, f17102, f17103, f17116, f17130, f17151, f17154, f17166, f17167, f17168, f17244, f17262, f17276, f17286, f17292, f17296, f17307, f17308, f17320, f17321, f17345, f17346, f17383, f17385, f17388, f17400, f17412, f17414, f17431, f17432, f17454, f17456, f17458, f17471, f17472, f17482, f17483, f17484, f17493, f17504, f17512, f17513, f17526, f17527, f17537, f17545, f17574, f17587, f17596, f17614, f17700, f17710, f17712, f17728, f17767, f17769, f17793, f17803, f17804, f17819, f17820, f17837, f17839, f17858, f17908, f17916, f17930, f17956, f17958, f17961, f17963, f17965, f17969, f17973, f17976, f17978, f17980, f17982, f17989, f18039, f18064, f18070, f18072, f18074, f18078, f18092, f18093, f18103, f18104, f18112, f18113, f18120, f18123, f18130, f18150, f18160, f18161, f18164, f18195, f18198, f18201, f18222, f18270, f18275, f18288, f18289, f18290, f18291, f18293, f18294, f18295, f18296, f18297, f18309, f18310, f18311, f18322, f18323, f18324, f18343, f18356, f18370, f18380, f18381, f18394, f18412, f18421, f18422, f18432, f18434, f18496, f18512, f18547, f18549, f18587, f18605, f18606, f18607, f18608, f18609, f18610, f18612, f18613, f18614, f18628, f18629, f18630, f18642, f18643, f18654, f18671, f18684, f18685, f18696, f18710, f18712, f18738, f18739, f18750, f18751, f18758, f18852, f18876, f18883, f18889, f18939, f18959, f18960, f18962, f18963, f18964, f18965, f18966, f18967, f18968, f18988, f18989, f18992, f19017, f19020, f19030, f19040, f19049, f19075, f19106, f19122, f19123, f19134, f19157, f19167, f19182, f19188, f19213, f19224, f19226, f19229, f19236, f19262, f19264, f19269, f19282, f19284, f19290, f19292, f19294, f19314, f19315, f19322, f19332, f19342, f19353, f19354, f19367])).
fof(f19367, plain, (~ spl24_23 | ~ spl24_112 | spl24_243 | ~ spl24_245), inference(avatar_contradiction_clause, [], [f19366])).
fof(f19366, plain, ($false | (~ spl24_23 | ~ spl24_112 | spl24_243 | ~ spl24_245)), inference(subsumption_resolution, [], [f19362, f1161])).
fof(f1161, plain, ((e1 = op(e0, e2)) | ~ spl24_112), inference(avatar_component_clause, [], [f1159])).
fof(f1159, plain, (spl24_112 <=> (e1 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_112])])).
fof(f19362, plain, (~ (e1 = op(e0, e2)) | (~ spl24_23 | spl24_243 | ~ spl24_245)), inference(backward_demodulation, [], [f19255, f787])).
fof(f787, plain, ((e2 = op(e4, e0)) | ~ spl24_23), inference(avatar_component_clause, [], [f785])).
fof(f785, plain, (spl24_23 <=> (e2 = op(e4, e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_23])])).
fof(f19255, plain, (~ (e1 = op(e0, op(e4, e0))) | (spl24_243 | ~ spl24_245)), inference(backward_demodulation, [], [f1846, f1853])).
fof(f1853, plain, ((e4 = op(e0, op(e0, e0))) | ~ spl24_245), inference(avatar_component_clause, [], [f1852])).
fof(f1852, plain, (spl24_245 <=> (e4 = op(e0, op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl24_245])])).
fof(f1846, plain, (~ (e1 = op(e0, op(op(e0, op(e0, e0)), e0))) | spl24_243), inference(avatar_component_clause, [], [f1844])).
fof(f1844, plain, (spl24_243 <=> (e1 = op(e0, op(op(e0, op(e0, e0)), e0)))), introduced(avatar_definition, [new_symbols(naming, [spl24_243])])).
fof(f19354, plain, (spl24_6 | ~ spl24_26 | ~ spl24_134), inference(avatar_split_clause, [], [f19348, f1303, f798, f714])).
fof(f714, plain, (spl24_6 <=> (e0 = op(e4, e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_6])])).
fof(f798, plain, (spl24_26 <=> (e0 = op(e3, e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_26])])).
fof(f1303, plain, (spl24_134 <=> (op(e3, e4) = op(e4, e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_134])])).
fof(f19348, plain, ((e0 = op(e4, e3)) | (~ spl24_26 | ~ spl24_134)), inference(backward_demodulation, [], [f1304, f800])).
fof(f800, plain, ((e0 = op(e3, e4)) | ~ spl24_26), inference(avatar_component_clause, [], [f798])).
fof(f1304, plain, ((op(e3, e4) = op(e4, e3)) | ~ spl24_134), inference(avatar_component_clause, [], [f1303])).
fof(f19353, plain, (~ spl24_1 | ~ spl24_26), inference(avatar_split_clause, [], [f19343, f798, f693])).
fof(f693, plain, (spl24_1 <=> (e0 = op(e4, e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_1])])).
fof(f19343, plain, (~ (e0 = op(e4, e4)) | ~ spl24_26), inference(backward_demodulation, [], [f431, f800])).
fof(f431, plain, ~ (op(e3, e4) = op(e4, e4)), inference(cnf_transformation, [], [f4])).
fof(f4, plain, (~ (op(e4, e3) = op(e4, e4)) & ~ (op(e4, e2) = op(e4, e4)) & ~ (op(e4, e2) = op(e4, e3)) & ~ (op(e4, e1) = op(e4, e4)) & ~ (op(e4, e1) = op(e4, e3)) & ~ (op(e4, e1) = op(e4, e2)) & ~ (op(e4, e0) = op(e4, e4)) & ~ (op(e4, e0) = op(e4, e3)) & ~ (op(e4, e0) = op(e4, e2)) & ~ (op(e4, e0) = op(e4, e1)) & ~ (op(e3, e3) = op(e3, e4)) & ~ (op(e3, e2) = op(e3, e4)) & ~ (op(e3, e2) = op(e3, e3)) & ~ (op(e3, e1) = op(e3, e4)) & ~ (op(e3, e1) = op(e3, e3)) & ~ (op(e3, e1) = op(e3, e2)) & ~ (op(e3, e0) = op(e3, e4)) & ~ (op(e3, e0) = op(e3, e3)) & ~ (op(e3, e0) = op(e3, e2)) & ~ (op(e3, e0) = op(e3, e1)) & ~ (op(e2, e3) = op(e2, e4)) & ~ (op(e2, e2) = op(e2, e4)) & ~ (op(e2, e2) = op(e2, e3)) & ~ (op(e2, e1) = op(e2, e4)) & ~ (op(e2, e1) = op(e2, e3)) & ~ (op(e2, e1) = op(e2, e2)) & ~ (op(e2, e0) = op(e2, e4)) & ~ (op(e2, e0) = op(e2, e3)) & ~ (op(e2, e0) = op(e2, e2)) & ~ (op(e2, e0) = op(e2, e1)) & ~ (op(e1, e3) = op(e1, e4)) & ~ (op(e1, e2) = op(e1, e4)) & ~ (op(e1, e2) = op(e1, e3)) & ~ (op(e1, e1) = op(e1, e4)) & ~ (op(e1, e1) = op(e1, e3)) & ~ (op(e1, e1) = op(e1, e2)) & ~ (op(e1, e0) = op(e1, e4)) & ~ (op(e1, e0) = op(e1, e3)) & ~ (op(e1, e0) = op(e1, e2)) & ~ (op(e1, e0) = op(e1, e1)) & ~ (op(e0, e3) = op(e0, e4)) & ~ (op(e0, e2) = op(e0, e4)) & ~ (op(e0, e2) = op(e0, e3)) & ~ (op(e0, e1) = op(e0, e4)) & ~ (op(e0, e1) = op(e0, e3)) & ~ (op(e0, e1) = op(e0, e2)) & ~ (op(e0, e0) = op(e0, e4)) & ~ (op(e0, e0) = op(e0, e3)) & ~ (op(e0, e0) = op(e0, e2)) & ~ (op(e0, e0) = op(e0, e1)) & ~ (op(e3, e4) = op(e4, e4)) & ~ (op(e2, e4) = op(e4, e4)) & ~ (op(e2, e4) = op(e3, e4)) & ~ (op(e1, e4) = op(e4, e4)) & ~ (op(e1, e4) = op(e3, e4)) & ~ (op(e1, e4) = op(e2, e4)) & ~ (op(e0, e4) = op(e4, e4)) & ~ (op(e0, e4) = op(e3, e4)) & ~ (op(e0, e4) = op(e2, e4)) & ~ (op(e0, e4) = op(e1, e4)) & ~ (op(e3, e3) = op(e4, e3)) & ~ (op(e2, e3) = op(e4, e3)) & ~ (op(e2, e3) = op(e3, e3)) & ~ (op(e1, e3) = op(e4, e3)) & ~ (op(e1, e3) = op(e3, e3)) & ~ (op(e1, e3) = op(e2, e3)) & ~ (op(e0, e3) = op(e4, e3)) & ~ (op(e0, e3) = op(e3, e3)) & ~ (op(e0, e3) = op(e2, e3)) & ~ (op(e0, e3) = op(e1, e3)) & ~ (op(e3, e2) = op(e4, e2)) & ~ (op(e2, e2) = op(e4, e2)) & ~ (op(e2, e2) = op(e3, e2)) & ~ (op(e1, e2) = op(e4, e2)) & ~ (op(e1, e2) = op(e3, e2)) & ~ (op(e1, e2) = op(e2, e2)) & ~ (op(e0, e2) = op(e4, e2)) & ~ (op(e0, e2) = op(e3, e2)) & ~ (op(e0, e2) = op(e2, e2)) & ~ (op(e0, e2) = op(e1, e2)) & ~ (op(e3, e1) = op(e4, e1)) & ~ (op(e2, e1) = op(e4, e1)) & ~ (op(e2, e1) = op(e3, e1)) & ~ (op(e1, e1) = op(e4, e1)) & ~ (op(e1, e1) = op(e3, e1)) & ~ (op(e1, e1) = op(e2, e1)) & ~ (op(e0, e1) = op(e4, e1)) & ~ (op(e0, e1) = op(e3, e1)) & ~ (op(e0, e1) = op(e2, e1)) & ~ (op(e0, e1) = op(e1, e1)) & ~ (op(e3, e0) = op(e4, e0)) & ~ (op(e2, e0) = op(e4, e0)) & ~ (op(e2, e0) = op(e3, e0)) & ~ (op(e1, e0) = op(e4, e0)) & ~ (op(e1, e0) = op(e3, e0)) & ~ (op(e1, e0) = op(e2, e0)) & ~ (op(e0, e0) = op(e4, e0)) & ~ (op(e0, e0) = op(e3, e0)) & ~ (op(e0, e0) = op(e2, e0)) & ~ (op(e0, e0) = op(e1, e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax4)).
fof(f19342, plain, (~ spl24_8 | ~ spl24_33), inference(avatar_split_clause, [], [f19334, f827, f722])).
fof(f722, plain, (spl24_8 <=> (e2 = op(e4, e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_8])])).
fof(f827, plain, (spl24_33 <=> (e2 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_33])])).
fof(f19334, plain, (~ (e2 = op(e4, e3)) | ~ spl24_33), inference(backward_demodulation, [], [f421, f829])).
fof(f829, plain, ((e2 = op(e3, e3)) | ~ spl24_33), inference(avatar_component_clause, [], [f827])).
fof(f421, plain, ~ (op(e3, e3) = op(e4, e3)), inference(cnf_transformation, [], [f4])).
fof(f19332, plain, (~ spl24_27 | ~ spl24_47), inference(avatar_split_clause, [], [f19326, f886, f802])).
fof(f802, plain, (spl24_27 <=> (e1 = op(e3, e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_27])])).
fof(f886, plain, (spl24_47 <=> (e1 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_47])])).
fof(f19326, plain, (~ (e1 = op(e3, e4)) | ~ spl24_47), inference(backward_demodulation, [], [f465, f888])).
fof(f888, plain, ((e1 = op(e3, e0)) | ~ spl24_47), inference(avatar_component_clause, [], [f886])).
fof(f465, plain, ~ (op(e3, e0) = op(e3, e4)), inference(cnf_transformation, [], [f4])).
fof(f19322, plain, (~ spl24_12 | ~ spl24_112), inference(avatar_split_clause, [], [f19316, f1159, f739])).
fof(f739, plain, (spl24_12 <=> (e1 = op(e4, e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_12])])).
fof(f19316, plain, (~ (e1 = op(e4, e2)) | ~ spl24_112), inference(backward_demodulation, [], [f405, f1161])).
fof(f405, plain, ~ (op(e0, e2) = op(e4, e2)), inference(cnf_transformation, [], [f4])).
fof(f19315, plain, (~ spl24_114 | ~ spl24_124), inference(avatar_split_clause, [], [f19308, f1209, f1167])).
fof(f1167, plain, (spl24_114 <=> (e3 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_114])])).
fof(f1209, plain, (spl24_124 <=> (op(e0, e0) = e3)), introduced(avatar_definition, [new_symbols(naming, [spl24_124])])).
fof(f19308, plain, (~ (e3 = op(e0, e2)) | ~ spl24_124), inference(backward_demodulation, [], [f433, f1211])).
fof(f1211, plain, ((op(e0, e0) = e3) | ~ spl24_124), inference(avatar_component_clause, [], [f1209])).
fof(f433, plain, ~ (op(e0, e0) = op(e0, e2)), inference(cnf_transformation, [], [f4])).
fof(f19314, plain, (~ spl24_24 | ~ spl24_124), inference(avatar_split_clause, [], [f19307, f1209, f789])).
fof(f789, plain, (spl24_24 <=> (e3 = op(e4, e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_24])])).
fof(f19307, plain, (~ (e3 = op(e4, e0)) | ~ spl24_124), inference(backward_demodulation, [], [f385, f1211])).
fof(f385, plain, ~ (op(e0, e0) = op(e4, e0)), inference(cnf_transformation, [], [f4])).
fof(f19294, plain, (~ spl24_8 | ~ spl24_40 | spl24_153), inference(avatar_split_clause, [], [f19293, f1395, f856, f722])).
fof(f856, plain, (spl24_40 <=> (e4 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_40])])).
fof(f1395, plain, (spl24_153 <=> (e2 = op(op(e3, e2), e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_153])])).
fof(f19293, plain, (~ (e2 = op(e4, e3)) | (~ spl24_40 | spl24_153)), inference(forward_demodulation, [], [f1397, f858])).
fof(f858, plain, ((e4 = op(e3, e2)) | ~ spl24_40), inference(avatar_component_clause, [], [f856])).
fof(f1397, plain, (~ (e2 = op(op(e3, e2), e3)) | spl24_153), inference(avatar_component_clause, [], [f1395])).
fof(f19292, plain, (spl24_14 | ~ spl24_40 | ~ spl24_154), inference(avatar_split_clause, [], [f19291, f1400, f856, f747])).
fof(f747, plain, (spl24_14 <=> (e3 = op(e4, e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_14])])).
fof(f1400, plain, (spl24_154 <=> (e3 = op(op(e3, e2), e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_154])])).
fof(f19291, plain, ((e3 = op(e4, e2)) | (~ spl24_40 | ~ spl24_154)), inference(forward_demodulation, [], [f1402, f858])).
fof(f1402, plain, ((e3 = op(op(e3, e2), e2)) | ~ spl24_154), inference(avatar_component_clause, [], [f1400])).
fof(f19290, plain, (spl24_23 | ~ spl24_75 | ~ spl24_178), inference(avatar_split_clause, [], [f19289, f1517, f1003, f785])).
fof(f1003, plain, (spl24_75 <=> (e4 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_75])])).
fof(f1517, plain, (spl24_178 <=> (e2 = op(op(e2, e0), e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_178])])).
fof(f19289, plain, ((e2 = op(e4, e0)) | (~ spl24_75 | ~ spl24_178)), inference(forward_demodulation, [], [f1519, f1005])).
fof(f1005, plain, ((e4 = op(e2, e0)) | ~ spl24_75), inference(avatar_component_clause, [], [f1003])).
fof(f1519, plain, ((e2 = op(op(e2, e0), e0)) | ~ spl24_178), inference(avatar_component_clause, [], [f1517])).
fof(f19284, plain, (~ spl24_24 | ~ spl24_110 | spl24_199), inference(avatar_split_clause, [], [f19283, f1621, f1150, f789])).
fof(f1150, plain, (spl24_110 <=> (e4 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_110])])).
fof(f1621, plain, (spl24_199 <=> (e3 = op(op(e0, e3), e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_199])])).
fof(f19283, plain, (~ (e3 = op(e4, e0)) | (~ spl24_110 | spl24_199)), inference(forward_demodulation, [], [f1623, f1152])).
fof(f1152, plain, ((e4 = op(e0, e3)) | ~ spl24_110), inference(avatar_component_clause, [], [f1150])).
fof(f1623, plain, (~ (e3 = op(op(e0, e3), e0)) | spl24_199), inference(avatar_component_clause, [], [f1621])).
fof(f19282, plain, (spl24_6 | ~ spl24_110 | ~ spl24_200), inference(avatar_split_clause, [], [f19281, f1626, f1150, f714])).
fof(f1626, plain, (spl24_200 <=> (e0 = op(op(e0, e3), e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_200])])).
fof(f19281, plain, ((e0 = op(e4, e3)) | (~ spl24_110 | ~ spl24_200)), inference(forward_demodulation, [], [f1628, f1152])).
fof(f1628, plain, ((e0 = op(op(e0, e3), e3)) | ~ spl24_200), inference(avatar_component_clause, [], [f1626])).
fof(f19269, plain, (~ spl24_47 | spl24_238 | ~ spl24_241), inference(avatar_split_clause, [], [f19268, f1834, f1816, f886])).
fof(f1816, plain, (spl24_238 <=> (e1 = op(e3, op(op(e3, op(e3, e3)), e3)))), introduced(avatar_definition, [new_symbols(naming, [spl24_238])])).
fof(f1834, plain, (spl24_241 <=> (e0 = op(op(e3, op(e3, e3)), e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_241])])).
fof(f19268, plain, (~ (e1 = op(e3, e0)) | (spl24_238 | ~ spl24_241)), inference(forward_demodulation, [], [f1818, f1835])).
fof(f1835, plain, ((e0 = op(op(e3, op(e3, e3)), e3)) | ~ spl24_241), inference(avatar_component_clause, [], [f1834])).
fof(f1818, plain, (~ (e1 = op(e3, op(op(e3, op(e3, e3)), e3))) | spl24_238), inference(avatar_component_clause, [], [f1816])).
fof(f19264, plain, (~ spl24_24 | spl24_244 | ~ spl24_245), inference(avatar_split_clause, [], [f19263, f1852, f1848, f789])).
fof(f1848, plain, (spl24_244 <=> (e3 = op(op(e0, op(e0, e0)), e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_244])])).
fof(f19263, plain, (~ (e3 = op(e4, e0)) | (spl24_244 | ~ spl24_245)), inference(forward_demodulation, [], [f1850, f1853])).
fof(f1850, plain, (~ (e3 = op(op(e0, op(e0, e0)), e0)) | spl24_244), inference(avatar_component_clause, [], [f1848])).
fof(f19262, plain, (~ spl24_23 | ~ spl24_245 | spl24_251), inference(avatar_split_clause, [], [f19257, f1882, f1852, f785])).
fof(f1882, plain, (spl24_251 <=> (e2 = op(op(e0, op(e0, e0)), e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_251])])).
fof(f19257, plain, (~ (e2 = op(e4, e0)) | (~ spl24_245 | spl24_251)), inference(backward_demodulation, [], [f1884, f1853])).
fof(f1884, plain, (~ (e2 = op(op(e0, op(e0, e0)), e0)) | spl24_251), inference(avatar_component_clause, [], [f1882])).
fof(f19236, plain, (~ spl24_115 | ~ spl24_75 | spl24_179), inference(avatar_split_clause, [], [f19235, f1522, f1003, f1171])).
fof(f1171, plain, (spl24_115 <=> (e4 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_115])])).
fof(f1522, plain, (spl24_179 <=> (op(e0, e2) = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_179])])).
fof(f19235, plain, (~ (e4 = op(e0, e2)) | (~ spl24_75 | spl24_179)), inference(forward_demodulation, [], [f1524, f1005])).
fof(f1524, plain, (~ (op(e0, e2) = op(e2, e0)) | spl24_179), inference(avatar_component_clause, [], [f1522])).
fof(f19229, plain, (~ spl24_15 | ~ spl24_20), inference(avatar_split_clause, [], [f19228, f772, f751])).
fof(f751, plain, (spl24_15 <=> (e4 = op(e4, e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_15])])).
fof(f772, plain, (spl24_20 <=> (e4 = op(e4, e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_20])])).
fof(f19228, plain, (~ (e4 = op(e4, e2)) | ~ spl24_20), inference(forward_demodulation, [], [f476, f774])).
fof(f774, plain, ((e4 = op(e4, e1)) | ~ spl24_20), inference(avatar_component_clause, [], [f772])).
fof(f476, plain, ~ (op(e4, e1) = op(e4, e2)), inference(cnf_transformation, [], [f4])).
fof(f19226, plain, (~ spl24_13 | ~ spl24_88), inference(avatar_split_clause, [], [f19225, f1058, f743])).
fof(f743, plain, (spl24_13 <=> (e2 = op(e4, e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_13])])).
fof(f1058, plain, (spl24_88 <=> (e2 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_88])])).
fof(f19225, plain, (~ (e2 = op(e4, e2)) | ~ spl24_88), inference(forward_demodulation, [], [f408, f1060])).
fof(f1060, plain, ((e2 = op(e1, e2)) | ~ spl24_88), inference(avatar_component_clause, [], [f1058])).
fof(f408, plain, ~ (op(e1, e2) = op(e4, e2)), inference(cnf_transformation, [], [f4])).
fof(f19224, plain, (~ spl24_15 | ~ spl24_40), inference(avatar_split_clause, [], [f19223, f856, f751])).
fof(f19223, plain, (~ (e4 = op(e4, e2)) | ~ spl24_40), inference(forward_demodulation, [], [f411, f858])).
fof(f411, plain, ~ (op(e3, e2) = op(e4, e2)), inference(cnf_transformation, [], [f4])).
fof(f19213, plain, (~ spl24_3 | ~ spl24_8), inference(avatar_split_clause, [], [f19208, f722, f701])).
fof(f701, plain, (spl24_3 <=> (e2 = op(e4, e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_3])])).
fof(f19208, plain, (~ (e2 = op(e4, e4)) | ~ spl24_8), inference(backward_demodulation, [], [f481, f724])).
fof(f724, plain, ((e2 = op(e4, e3)) | ~ spl24_8), inference(avatar_component_clause, [], [f722])).
fof(f481, plain, ~ (op(e4, e3) = op(e4, e4)), inference(cnf_transformation, [], [f4])).
fof(f19188, plain, (~ spl24_30 | ~ spl24_40), inference(avatar_split_clause, [], [f19185, f856, f814])).
fof(f814, plain, (spl24_30 <=> (e4 = op(e3, e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_30])])).
fof(f19185, plain, (~ (e4 = op(e3, e4)) | ~ spl24_40), inference(backward_demodulation, [], [f470, f858])).
fof(f470, plain, ~ (op(e3, e2) = op(e3, e4)), inference(cnf_transformation, [], [f4])).
fof(f19182, plain, (~ spl24_38 | ~ spl24_48), inference(avatar_split_clause, [], [f19175, f890, f848])).
fof(f848, plain, (spl24_38 <=> (e2 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_38])])).
fof(f890, plain, (spl24_48 <=> (e2 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_48])])).
fof(f19175, plain, (~ (e2 = op(e3, e2)) | ~ spl24_48), inference(backward_demodulation, [], [f463, f892])).
fof(f892, plain, ((e2 = op(e3, e0)) | ~ spl24_48), inference(avatar_component_clause, [], [f890])).
fof(f463, plain, ~ (op(e3, e0) = op(e3, e2)), inference(cnf_transformation, [], [f4])).
fof(f19167, plain, (~ spl24_32 | ~ spl24_57), inference(avatar_split_clause, [], [f19158, f928, f823])).
fof(f823, plain, (spl24_32 <=> (e1 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_32])])).
fof(f928, plain, (spl24_57 <=> (e1 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_57])])).
fof(f19158, plain, (~ (e1 = op(e3, e3)) | ~ spl24_57), inference(backward_demodulation, [], [f419, f930])).
fof(f930, plain, ((e1 = op(e2, e3)) | ~ spl24_57), inference(avatar_component_clause, [], [f928])).
fof(f419, plain, ~ (op(e2, e3) = op(e3, e3)), inference(cnf_transformation, [], [f4])).
fof(f19157, plain, (~ spl24_51 | ~ spl24_61), inference(avatar_split_clause, [], [f19148, f945, f903])).
fof(f903, plain, (spl24_51 <=> (e0 = op(e2, e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_51])])).
fof(f945, plain, (spl24_61 <=> (e0 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_61])])).
fof(f19148, plain, (~ (e0 = op(e2, e4)) | ~ spl24_61), inference(backward_demodulation, [], [f460, f947])).
fof(f947, plain, ((e0 = op(e2, e2)) | ~ spl24_61), inference(avatar_component_clause, [], [f945])).
fof(f460, plain, ~ (op(e2, e2) = op(e2, e4)), inference(cnf_transformation, [], [f4])).
fof(f19134, plain, (~ spl24_30 | ~ spl24_80), inference(avatar_split_clause, [], [f19128, f1024, f814])).
fof(f1024, plain, (spl24_80 <=> (e4 = op(e1, e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_80])])).
fof(f19128, plain, (~ (e4 = op(e3, e4)) | ~ spl24_80), inference(backward_demodulation, [], [f427, f1026])).
fof(f1026, plain, ((e4 = op(e1, e4)) | ~ spl24_80), inference(avatar_component_clause, [], [f1024])).
fof(f427, plain, ~ (op(e1, e4) = op(e3, e4)), inference(cnf_transformation, [], [f4])).
fof(f19123, plain, (~ spl24_79 | ~ spl24_84), inference(avatar_split_clause, [], [f19113, f1041, f1020])).
fof(f1020, plain, (spl24_79 <=> (e3 = op(e1, e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_79])])).
fof(f1041, plain, (spl24_84 <=> (e3 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_84])])).
fof(f19113, plain, (~ (e3 = op(e1, e4)) | ~ spl24_84), inference(backward_demodulation, [], [f451, f1043])).
fof(f1043, plain, ((e3 = op(e1, e3)) | ~ spl24_84), inference(avatar_component_clause, [], [f1041])).
fof(f451, plain, ~ (op(e1, e3) = op(e1, e4)), inference(cnf_transformation, [], [f4])).
fof(f19122, plain, (~ spl24_34 | ~ spl24_84), inference(avatar_split_clause, [], [f19111, f1041, f831])).
fof(f831, plain, (spl24_34 <=> (e3 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_34])])).
fof(f19111, plain, (~ (e3 = op(e3, e3)) | ~ spl24_84), inference(backward_demodulation, [], [f417, f1043])).
fof(f417, plain, ~ (op(e1, e3) = op(e3, e3)), inference(cnf_transformation, [], [f4])).
fof(f19106, plain, (~ spl24_38 | ~ spl24_88), inference(avatar_split_clause, [], [f19094, f1058, f848])).
fof(f19094, plain, (~ (e2 = op(e3, e2)) | ~ spl24_88), inference(backward_demodulation, [], [f407, f1060])).
fof(f407, plain, ~ (op(e1, e2) = op(e3, e2)), inference(cnf_transformation, [], [f4])).
fof(f19075, plain, (~ spl24_82 | ~ spl24_92), inference(avatar_split_clause, [], [f19060, f1075, f1033])).
fof(f1033, plain, (spl24_82 <=> (e1 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_82])])).
fof(f1075, plain, (spl24_92 <=> (e1 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_92])])).
fof(f19060, plain, (~ (e1 = op(e1, e3)) | ~ spl24_92), inference(backward_demodulation, [], [f447, f1077])).
fof(f1077, plain, ((e1 = op(e1, e1)) | ~ spl24_92), inference(avatar_component_clause, [], [f1075])).
fof(f447, plain, ~ (op(e1, e1) = op(e1, e3)), inference(cnf_transformation, [], [f4])).
fof(f19049, plain, (~ spl24_3 | ~ spl24_103), inference(avatar_split_clause, [], [f19044, f1121, f701])).
fof(f1121, plain, (spl24_103 <=> (e2 = op(e0, e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_103])])).
fof(f19044, plain, (~ (e2 = op(e4, e4)) | ~ spl24_103), inference(backward_demodulation, [], [f425, f1123])).
fof(f1123, plain, ((e2 = op(e0, e4)) | ~ spl24_103), inference(avatar_component_clause, [], [f1121])).
fof(f425, plain, ~ (op(e0, e4) = op(e4, e4)), inference(cnf_transformation, [], [f4])).
fof(f19040, plain, (~ spl24_10 | ~ spl24_110), inference(avatar_split_clause, [], [f19035, f1150, f730])).
fof(f730, plain, (spl24_10 <=> (e4 = op(e4, e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_10])])).
fof(f19035, plain, (~ (e4 = op(e4, e3)) | ~ spl24_110), inference(backward_demodulation, [], [f415, f1152])).
fof(f415, plain, ~ (op(e0, e3) = op(e4, e3)), inference(cnf_transformation, [], [f4])).
fof(f19030, plain, (~ spl24_64 | ~ spl24_114), inference(avatar_split_clause, [], [f19022, f1167, f957])).
fof(f957, plain, (spl24_64 <=> (e3 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_64])])).
fof(f19022, plain, (~ (e3 = op(e2, e2)) | ~ spl24_114), inference(backward_demodulation, [], [f403, f1169])).
fof(f1169, plain, ((e3 = op(e0, e2)) | ~ spl24_114), inference(avatar_component_clause, [], [f1167])).
fof(f403, plain, ~ (op(e0, e2) = op(e2, e2)), inference(cnf_transformation, [], [f4])).
fof(f19020, plain, (~ spl24_86 | ~ spl24_116 | ~ spl24_194), inference(avatar_split_clause, [], [f19013, f1597, f1176, f1050])).
fof(f1050, plain, (spl24_86 <=> (e0 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_86])])).
fof(f1176, plain, (spl24_116 <=> (e0 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_116])])).
fof(f1597, plain, (spl24_194 <=> (op(e0, e1) = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_194])])).
fof(f19013, plain, (~ (e0 = op(e1, e2)) | (~ spl24_116 | ~ spl24_194)), inference(backward_demodulation, [], [f18922, f1178])).
fof(f1178, plain, ((e0 = op(e0, e1)) | ~ spl24_116), inference(avatar_component_clause, [], [f1176])).
fof(f18922, plain, (~ (op(e0, e1) = op(e1, e2)) | ~ spl24_194), inference(backward_demodulation, [], [f443, f1598])).
fof(f1598, plain, ((op(e0, e1) = op(e1, e0)) | ~ spl24_194), inference(avatar_component_clause, [], [f1597])).
fof(f443, plain, ~ (op(e1, e0) = op(e1, e2)), inference(cnf_transformation, [], [f4])).
fof(f19017, plain, (~ spl24_106 | ~ spl24_116), inference(avatar_split_clause, [], [f19002, f1176, f1134])).
fof(f1134, plain, (spl24_106 <=> (e0 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_106])])).
fof(f19002, plain, (~ (e0 = op(e0, e3)) | ~ spl24_116), inference(backward_demodulation, [], [f437, f1178])).
fof(f437, plain, ~ (op(e0, e1) = op(e0, e3)), inference(cnf_transformation, [], [f4])).
fof(f18992, plain, (~ spl24_119 | ~ spl24_122 | spl24_248), inference(avatar_split_clause, [], [f18982, f1866, f1201, f1188])).
fof(f1188, plain, (spl24_119 <=> (e3 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_119])])).
fof(f1201, plain, (spl24_122 <=> (op(e0, e0) = e1)), introduced(avatar_definition, [new_symbols(naming, [spl24_122])])).
fof(f1866, plain, (spl24_248 <=> (e3 = op(e0, op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl24_248])])).
fof(f18982, plain, (~ (e3 = op(e0, e1)) | (~ spl24_122 | spl24_248)), inference(backward_demodulation, [], [f1868, f1203])).
fof(f1203, plain, ((op(e0, e0) = e1) | ~ spl24_122), inference(avatar_component_clause, [], [f1201])).
fof(f1868, plain, (~ (e3 = op(e0, op(e0, e0))) | spl24_248), inference(avatar_component_clause, [], [f1866])).
fof(f18989, plain, (~ spl24_102 | ~ spl24_122), inference(avatar_split_clause, [], [f18975, f1201, f1117])).
fof(f1117, plain, (spl24_102 <=> (e1 = op(e0, e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_102])])).
fof(f18975, plain, (~ (e1 = op(e0, e4)) | ~ spl24_122), inference(backward_demodulation, [], [f435, f1203])).
fof(f435, plain, ~ (op(e0, e0) = op(e0, e4)), inference(cnf_transformation, [], [f4])).
fof(f18988, plain, (~ spl24_72 | ~ spl24_122), inference(avatar_split_clause, [], [f18970, f1201, f991])).
fof(f991, plain, (spl24_72 <=> (e1 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_72])])).
fof(f18970, plain, (~ (e1 = op(e2, e0)) | ~ spl24_122), inference(backward_demodulation, [], [f383, f1203])).
fof(f383, plain, ~ (op(e0, e0) = op(e2, e0)), inference(cnf_transformation, [], [f4])).
fof(f18968, plain, (spl24_20 | ~ spl24_127), inference(avatar_split_clause, [], [f18955, f1222, f772])).
fof(f1222, plain, (spl24_127 <=> (e1 = unit)), introduced(avatar_definition, [new_symbols(naming, [spl24_127])])).
fof(f18955, plain, ((e4 = op(e4, e1)) | ~ spl24_127), inference(backward_demodulation, [], [f330, f1224])).
fof(f1224, plain, ((e1 = unit) | ~ spl24_127), inference(avatar_component_clause, [], [f1222])).
fof(f330, plain, (e4 = op(e4, unit)), inference(cnf_transformation, [], [f2])).
fof(f2, plain, (((e4 = unit) | (e3 = unit) | (e2 = unit) | (e1 = unit) | (e0 = unit)) & (e4 = op(e4, unit)) & (e4 = op(unit, e4)) & (e3 = op(e3, unit)) & (e3 = op(unit, e3)) & (e2 = op(e2, unit)) & (e2 = op(unit, e2)) & (e1 = op(e1, unit)) & (e1 = op(unit, e1)) & (e0 = op(e0, unit)) & (e0 = op(unit, e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax2)).
fof(f18967, plain, (spl24_80 | ~ spl24_127), inference(avatar_split_clause, [], [f18954, f1222, f1024])).
fof(f18954, plain, ((e4 = op(e1, e4)) | ~ spl24_127), inference(backward_demodulation, [], [f329, f1224])).
fof(f329, plain, (e4 = op(unit, e4)), inference(cnf_transformation, [], [f2])).
fof(f18966, plain, (spl24_44 | ~ spl24_127), inference(avatar_split_clause, [], [f18953, f1222, f873])).
fof(f873, plain, (spl24_44 <=> (e3 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_44])])).
fof(f18953, plain, ((e3 = op(e3, e1)) | ~ spl24_127), inference(backward_demodulation, [], [f328, f1224])).
fof(f328, plain, (e3 = op(e3, unit)), inference(cnf_transformation, [], [f2])).
fof(f18965, plain, (spl24_84 | ~ spl24_127), inference(avatar_split_clause, [], [f18952, f1222, f1041])).
fof(f18952, plain, ((e3 = op(e1, e3)) | ~ spl24_127), inference(backward_demodulation, [], [f327, f1224])).
fof(f327, plain, (e3 = op(unit, e3)), inference(cnf_transformation, [], [f2])).
fof(f18964, plain, (spl24_68 | ~ spl24_127), inference(avatar_split_clause, [], [f18951, f1222, f974])).
fof(f974, plain, (spl24_68 <=> (e2 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_68])])).
fof(f18951, plain, ((e2 = op(e2, e1)) | ~ spl24_127), inference(backward_demodulation, [], [f326, f1224])).
fof(f326, plain, (e2 = op(e2, unit)), inference(cnf_transformation, [], [f2])).
fof(f18963, plain, (spl24_88 | ~ spl24_127), inference(avatar_split_clause, [], [f18950, f1222, f1058])).
fof(f18950, plain, ((e2 = op(e1, e2)) | ~ spl24_127), inference(backward_demodulation, [], [f325, f1224])).
fof(f325, plain, (e2 = op(unit, e2)), inference(cnf_transformation, [], [f2])).
fof(f18962, plain, (spl24_92 | ~ spl24_127), inference(avatar_split_clause, [], [f18949, f1222, f1075])).
fof(f18949, plain, ((e1 = op(e1, e1)) | ~ spl24_127), inference(backward_demodulation, [], [f324, f1224])).
fof(f324, plain, (e1 = op(e1, unit)), inference(cnf_transformation, [], [f2])).
fof(f18960, plain, (spl24_116 | ~ spl24_127), inference(avatar_split_clause, [], [f18947, f1222, f1176])).
fof(f18947, plain, ((e0 = op(e0, e1)) | ~ spl24_127), inference(backward_demodulation, [], [f322, f1224])).
fof(f322, plain, (e0 = op(e0, unit)), inference(cnf_transformation, [], [f2])).
fof(f18959, plain, (spl24_96 | ~ spl24_127), inference(avatar_split_clause, [], [f18946, f1222, f1092])).
fof(f1092, plain, (spl24_96 <=> (e0 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_96])])).
fof(f18946, plain, ((e0 = op(e1, e0)) | ~ spl24_127), inference(backward_demodulation, [], [f321, f1224])).
fof(f321, plain, (e0 = op(unit, e0)), inference(cnf_transformation, [], [f2])).
fof(f18939, plain, (~ spl24_24 | spl24_50 | ~ spl24_145), inference(avatar_contradiction_clause, [], [f18938])).
fof(f18938, plain, ($false | (~ spl24_24 | spl24_50 | ~ spl24_145)), inference(subsumption_resolution, [], [f18937, f899])).
fof(f899, plain, (~ (e4 = op(e3, e0)) | spl24_50), inference(avatar_component_clause, [], [f898])).
fof(f898, plain, (spl24_50 <=> (e4 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_50])])).
fof(f18937, plain, ((e4 = op(e3, e0)) | (~ spl24_24 | ~ spl24_145)), inference(forward_demodulation, [], [f1357, f791])).
fof(f791, plain, ((e3 = op(e4, e0)) | ~ spl24_24), inference(avatar_component_clause, [], [f789])).
fof(f1357, plain, ((e4 = op(op(e4, e0), e0)) | ~ spl24_145), inference(avatar_component_clause, [], [f1355])).
fof(f1355, plain, (spl24_145 <=> (e4 = op(op(e4, e0), e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_145])])).
fof(f18889, plain, (~ spl24_104 | ~ spl24_24 | spl24_146), inference(avatar_split_clause, [], [f18888, f1360, f789, f1125])).
fof(f1125, plain, (spl24_104 <=> (e3 = op(e0, e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_104])])).
fof(f1360, plain, (spl24_146 <=> (op(e0, e4) = op(e4, e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_146])])).
fof(f18888, plain, (~ (e3 = op(e0, e4)) | (~ spl24_24 | spl24_146)), inference(forward_demodulation, [], [f1362, f791])).
fof(f1362, plain, (~ (op(e0, e4) = op(e4, e0)) | spl24_146), inference(avatar_component_clause, [], [f1360])).
fof(f18883, plain, (~ spl24_52 | ~ spl24_12 | spl24_138), inference(avatar_split_clause, [], [f18882, f1322, f739, f907])).
fof(f907, plain, (spl24_52 <=> (e1 = op(e2, e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_52])])).
fof(f1322, plain, (spl24_138 <=> (op(e2, e4) = op(e4, e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_138])])).
fof(f18882, plain, (~ (e1 = op(e2, e4)) | (~ spl24_12 | spl24_138)), inference(forward_demodulation, [], [f1324, f741])).
fof(f741, plain, ((e1 = op(e4, e2)) | ~ spl24_12), inference(avatar_component_clause, [], [f739])).
fof(f1324, plain, (~ (op(e2, e4) = op(e4, e2)) | spl24_138), inference(avatar_component_clause, [], [f1322])).
fof(f18876, plain, (~ spl24_12 | ~ spl24_86 | ~ spl24_137), inference(avatar_contradiction_clause, [], [f18875])).
fof(f18875, plain, ($false | (~ spl24_12 | ~ spl24_86 | ~ spl24_137)), inference(subsumption_resolution, [], [f18874, f485])).
fof(f485, plain, ~ (e0 = e4), inference(cnf_transformation, [], [f5])).
fof(f5, plain, (~ (e3 = e4) & ~ (e2 = e4) & ~ (e2 = e3) & ~ (e1 = e4) & ~ (e1 = e3) & ~ (e1 = e2) & ~ (e0 = e4) & ~ (e0 = e3) & ~ (e0 = e2) & ~ (e0 = e1)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax5)).
fof(f18874, plain, ((e0 = e4) | (~ spl24_12 | ~ spl24_86 | ~ spl24_137)), inference(forward_demodulation, [], [f18873, f1052])).
fof(f1052, plain, ((e0 = op(e1, e2)) | ~ spl24_86), inference(avatar_component_clause, [], [f1050])).
fof(f18873, plain, ((e4 = op(e1, e2)) | (~ spl24_12 | ~ spl24_137)), inference(backward_demodulation, [], [f1319, f741])).
fof(f1319, plain, ((e4 = op(op(e4, e2), e2)) | ~ spl24_137), inference(avatar_component_clause, [], [f1317])).
fof(f1317, plain, (spl24_137 <=> (e4 = op(op(e4, e2), e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_137])])).
fof(f18852, plain, (~ spl24_4 | ~ spl24_24), inference(avatar_split_clause, [], [f18846, f789, f705])).
fof(f705, plain, (spl24_4 <=> (e3 = op(e4, e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_4])])).
fof(f18846, plain, (~ (e3 = op(e4, e4)) | ~ spl24_24), inference(backward_demodulation, [], [f475, f791])).
fof(f475, plain, ~ (op(e4, e0) = op(e4, e4)), inference(cnf_transformation, [], [f4])).
fof(f18758, plain, (~ spl24_55 | ~ spl24_70), inference(avatar_split_clause, [], [f18755, f982, f919])).
fof(f919, plain, (spl24_55 <=> (e4 = op(e2, e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_55])])).
fof(f982, plain, (spl24_70 <=> (e4 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_70])])).
fof(f18755, plain, (~ (e4 = op(e2, e4)) | ~ spl24_70), inference(backward_demodulation, [], [f458, f984])).
fof(f984, plain, ((e4 = op(e2, e1)) | ~ spl24_70), inference(avatar_component_clause, [], [f982])).
fof(f458, plain, ~ (op(e2, e1) = op(e2, e4)), inference(cnf_transformation, [], [f4])).
fof(f18751, plain, (~ spl24_67 | ~ spl24_72), inference(avatar_split_clause, [], [f18741, f991, f970])).
fof(f970, plain, (spl24_67 <=> (e1 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_67])])).
fof(f18741, plain, (~ (e1 = op(e2, e1)) | ~ spl24_72), inference(backward_demodulation, [], [f452, f993])).
fof(f993, plain, ((e1 = op(e2, e0)) | ~ spl24_72), inference(avatar_component_clause, [], [f991])).
fof(f452, plain, ~ (op(e2, e0) = op(e2, e1)), inference(cnf_transformation, [], [f4])).
fof(f18750, plain, (~ spl24_22 | ~ spl24_72), inference(avatar_split_clause, [], [f18740, f991, f781])).
fof(f781, plain, (spl24_22 <=> (e1 = op(e4, e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_22])])).
fof(f18740, plain, (~ (e1 = op(e4, e0)) | ~ spl24_72), inference(backward_demodulation, [], [f390, f993])).
fof(f390, plain, ~ (op(e2, e0) = op(e4, e0)), inference(cnf_transformation, [], [f4])).
fof(f18739, plain, (~ spl24_6 | ~ spl24_79 | spl24_213 | ~ spl24_230), inference(avatar_split_clause, [], [f18737, f1769, f1693, f1020, f714])).
fof(f1693, plain, (spl24_213 <=> (e0 = op(e4, op(op(e4, op(e4, e4)), e4)))), introduced(avatar_definition, [new_symbols(naming, [spl24_213])])).
fof(f1769, plain, (spl24_230 <=> (e1 = op(e4, op(e4, e4)))), introduced(avatar_definition, [new_symbols(naming, [spl24_230])])).
fof(f18737, plain, (~ (e0 = op(e4, e3)) | (~ spl24_79 | spl24_213 | ~ spl24_230)), inference(backward_demodulation, [], [f18564, f1022])).
fof(f1022, plain, ((e3 = op(e1, e4)) | ~ spl24_79), inference(avatar_component_clause, [], [f1020])).
fof(f18564, plain, (~ (e0 = op(e4, op(e1, e4))) | (spl24_213 | ~ spl24_230)), inference(forward_demodulation, [], [f1695, f1770])).
fof(f1770, plain, ((e1 = op(e4, op(e4, e4))) | ~ spl24_230), inference(avatar_component_clause, [], [f1769])).
fof(f1695, plain, (~ (e0 = op(e4, op(op(e4, op(e4, e4)), e4))) | spl24_213), inference(avatar_component_clause, [], [f1693])).
fof(f18738, plain, (~ spl24_4 | ~ spl24_79), inference(avatar_split_clause, [], [f18732, f1020, f705])).
fof(f18732, plain, (~ (e3 = op(e4, e4)) | ~ spl24_79), inference(backward_demodulation, [], [f428, f1022])).
fof(f428, plain, ~ (op(e1, e4) = op(e4, e4)), inference(cnf_transformation, [], [f4])).
fof(f18712, plain, (~ spl24_86 | spl24_118 | ~ spl24_187), inference(avatar_contradiction_clause, [], [f18711])).
fof(f18711, plain, ($false | (~ spl24_86 | spl24_118 | ~ spl24_187)), inference(subsumption_resolution, [], [f18708, f1185])).
fof(f1185, plain, (~ (e2 = op(e0, e1)) | spl24_118), inference(avatar_component_clause, [], [f1184])).
fof(f1184, plain, (spl24_118 <=> (e2 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_118])])).
fof(f18708, plain, ((e2 = op(e0, e1)) | (~ spl24_86 | ~ spl24_187)), inference(backward_demodulation, [], [f1562, f1052])).
fof(f1562, plain, ((e2 = op(op(e1, e2), e1)) | ~ spl24_187), inference(avatar_component_clause, [], [f1561])).
fof(f1561, plain, (spl24_187 <=> (e2 = op(op(e1, e2), e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_187])])).
fof(f18710, plain, (~ spl24_76 | ~ spl24_86), inference(avatar_split_clause, [], [f18702, f1050, f1008])).
fof(f1008, plain, (spl24_76 <=> (e0 = op(e1, e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_76])])).
fof(f18702, plain, (~ (e0 = op(e1, e4)) | ~ spl24_86), inference(backward_demodulation, [], [f450, f1052])).
fof(f450, plain, ~ (op(e1, e2) = op(e1, e4)), inference(cnf_transformation, [], [f4])).
fof(f18696, plain, (~ spl24_18 | ~ spl24_93), inference(avatar_split_clause, [], [f18689, f1079, f764])).
fof(f764, plain, (spl24_18 <=> (e2 = op(e4, e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_18])])).
fof(f1079, plain, (spl24_93 <=> (e2 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_93])])).
fof(f18689, plain, (~ (e2 = op(e4, e1)) | ~ spl24_93), inference(backward_demodulation, [], [f398, f1081])).
fof(f1081, plain, ((e2 = op(e1, e1)) | ~ spl24_93), inference(avatar_component_clause, [], [f1079])).
fof(f398, plain, ~ (op(e1, e1) = op(e4, e1)), inference(cnf_transformation, [], [f4])).
fof(f18685, plain, (spl24_16 | ~ spl24_100 | ~ spl24_192), inference(avatar_split_clause, [], [f18682, f1587, f1108, f756])).
fof(f756, plain, (spl24_16 <=> (e0 = op(e4, e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_16])])).
fof(f1108, plain, (spl24_100 <=> (e4 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_100])])).
fof(f1587, plain, (spl24_192 <=> (e0 = op(op(e1, e0), e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_192])])).
fof(f18682, plain, ((e0 = op(e4, e1)) | (~ spl24_100 | ~ spl24_192)), inference(backward_demodulation, [], [f1588, f1110])).
fof(f1110, plain, ((e4 = op(e1, e0)) | ~ spl24_100), inference(avatar_component_clause, [], [f1108])).
fof(f1588, plain, ((e0 = op(op(e1, e0), e1)) | ~ spl24_192), inference(avatar_component_clause, [], [f1587])).
fof(f18684, plain, (~ spl24_85 | ~ spl24_100), inference(avatar_split_clause, [], [f18680, f1108, f1045])).
fof(f1045, plain, (spl24_85 <=> (e4 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_85])])).
fof(f18680, plain, (~ (e4 = op(e1, e3)) | ~ spl24_100), inference(backward_demodulation, [], [f444, f1110])).
fof(f444, plain, ~ (op(e1, e0) = op(e1, e3)), inference(cnf_transformation, [], [f4])).
fof(f18671, plain, (~ spl24_6 | ~ spl24_106), inference(avatar_split_clause, [], [f18658, f1134, f714])).
fof(f18658, plain, (~ (e0 = op(e4, e3)) | ~ spl24_106), inference(backward_demodulation, [], [f415, f1136])).
fof(f1136, plain, ((e0 = op(e0, e3)) | ~ spl24_106), inference(avatar_component_clause, [], [f1134])).
fof(f18654, plain, (~ spl24_15 | ~ spl24_115), inference(avatar_split_clause, [], [f18649, f1171, f751])).
fof(f18649, plain, (~ (e4 = op(e4, e2)) | ~ spl24_115), inference(backward_demodulation, [], [f405, f1173])).
fof(f1173, plain, ((e4 = op(e0, e2)) | ~ spl24_115), inference(avatar_component_clause, [], [f1171])).
fof(f18643, plain, (~ spl24_41 | ~ spl24_119 | spl24_206), inference(avatar_split_clause, [], [f18640, f1656, f1188, f861])).
fof(f861, plain, (spl24_41 <=> (e0 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_41])])).
fof(f1656, plain, (spl24_206 <=> (e0 = op(op(e0, e1), e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_206])])).
fof(f18640, plain, (~ (e0 = op(e3, e1)) | (~ spl24_119 | spl24_206)), inference(backward_demodulation, [], [f1657, f1190])).
fof(f1190, plain, ((e3 = op(e0, e1)) | ~ spl24_119), inference(avatar_component_clause, [], [f1188])).
fof(f1657, plain, (~ (e0 = op(op(e0, e1), e1)) | spl24_206), inference(avatar_component_clause, [], [f1656])).
fof(f18642, plain, (~ spl24_94 | ~ spl24_119), inference(avatar_split_clause, [], [f18631, f1188, f1083])).
fof(f1083, plain, (spl24_94 <=> (e3 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_94])])).
fof(f18631, plain, (~ (e3 = op(e1, e1)) | ~ spl24_119), inference(backward_demodulation, [], [f392, f1190])).
fof(f392, plain, ~ (op(e0, e1) = op(e1, e1)), inference(cnf_transformation, [], [f4])).
fof(f18630, plain, (spl24_115 | ~ spl24_123 | ~ spl24_245), inference(avatar_split_clause, [], [f18624, f1852, f1205, f1171])).
fof(f1205, plain, (spl24_123 <=> (op(e0, e0) = e2)), introduced(avatar_definition, [new_symbols(naming, [spl24_123])])).
fof(f18624, plain, ((e4 = op(e0, e2)) | (~ spl24_123 | ~ spl24_245)), inference(backward_demodulation, [], [f1853, f1207])).
fof(f1207, plain, ((op(e0, e0) = e2) | ~ spl24_123), inference(avatar_component_clause, [], [f1205])).
fof(f18629, plain, (~ spl24_108 | ~ spl24_123), inference(avatar_split_clause, [], [f18621, f1205, f1142])).
fof(f1142, plain, (spl24_108 <=> (e2 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_108])])).
fof(f18621, plain, (~ (e2 = op(e0, e3)) | ~ spl24_123), inference(backward_demodulation, [], [f434, f1207])).
fof(f434, plain, ~ (op(e0, e0) = op(e0, e3)), inference(cnf_transformation, [], [f4])).
fof(f18628, plain, (~ spl24_98 | ~ spl24_123), inference(avatar_split_clause, [], [f18615, f1205, f1100])).
fof(f1100, plain, (spl24_98 <=> (e2 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_98])])).
fof(f18615, plain, (~ (e2 = op(e1, e0)) | ~ spl24_123), inference(backward_demodulation, [], [f382, f1207])).
fof(f382, plain, ~ (op(e0, e0) = op(e1, e0)), inference(cnf_transformation, [], [f4])).
fof(f18614, plain, (spl24_10 | ~ spl24_129), inference(avatar_split_clause, [], [f18603, f1230, f730])).
fof(f1230, plain, (spl24_129 <=> (e3 = unit)), introduced(avatar_definition, [new_symbols(naming, [spl24_129])])).
fof(f18603, plain, ((e4 = op(e4, e3)) | ~ spl24_129), inference(backward_demodulation, [], [f330, f1232])).
fof(f1232, plain, ((e3 = unit) | ~ spl24_129), inference(avatar_component_clause, [], [f1230])).
fof(f18613, plain, (spl24_30 | ~ spl24_129), inference(avatar_split_clause, [], [f18602, f1230, f814])).
fof(f18602, plain, ((e4 = op(e3, e4)) | ~ spl24_129), inference(backward_demodulation, [], [f329, f1232])).
fof(f18612, plain, (spl24_34 | ~ spl24_129), inference(avatar_split_clause, [], [f18601, f1230, f831])).
fof(f18601, plain, ((e3 = op(e3, e3)) | ~ spl24_129), inference(backward_demodulation, [], [f328, f1232])).
fof(f18610, plain, (spl24_58 | ~ spl24_129), inference(avatar_split_clause, [], [f18599, f1230, f932])).
fof(f932, plain, (spl24_58 <=> (e2 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_58])])).
fof(f18599, plain, ((e2 = op(e2, e3)) | ~ spl24_129), inference(backward_demodulation, [], [f326, f1232])).
fof(f18609, plain, (spl24_38 | ~ spl24_129), inference(avatar_split_clause, [], [f18598, f1230, f848])).
fof(f18598, plain, ((e2 = op(e3, e2)) | ~ spl24_129), inference(backward_demodulation, [], [f325, f1232])).
fof(f18608, plain, (spl24_82 | ~ spl24_129), inference(avatar_split_clause, [], [f18597, f1230, f1033])).
fof(f18597, plain, ((e1 = op(e1, e3)) | ~ spl24_129), inference(backward_demodulation, [], [f324, f1232])).
fof(f18607, plain, (spl24_42 | ~ spl24_129), inference(avatar_split_clause, [], [f18596, f1230, f865])).
fof(f865, plain, (spl24_42 <=> (e1 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_42])])).
fof(f18596, plain, ((e1 = op(e3, e1)) | ~ spl24_129), inference(backward_demodulation, [], [f323, f1232])).
fof(f323, plain, (e1 = op(unit, e1)), inference(cnf_transformation, [], [f2])).
fof(f18606, plain, (spl24_106 | ~ spl24_129), inference(avatar_split_clause, [], [f18595, f1230, f1134])).
fof(f18595, plain, ((e0 = op(e0, e3)) | ~ spl24_129), inference(backward_demodulation, [], [f322, f1232])).
fof(f18605, plain, (spl24_46 | ~ spl24_129), inference(avatar_split_clause, [], [f18594, f1230, f882])).
fof(f882, plain, (spl24_46 <=> (e0 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_46])])).
fof(f18594, plain, ((e0 = op(e3, e0)) | ~ spl24_129), inference(backward_demodulation, [], [f321, f1232])).
fof(f18587, plain, (~ spl24_22 | ~ spl24_102 | spl24_146), inference(avatar_split_clause, [], [f18586, f1360, f1117, f781])).
fof(f18586, plain, (~ (e1 = op(e4, e0)) | (~ spl24_102 | spl24_146)), inference(forward_demodulation, [], [f1362, f1119])).
fof(f1119, plain, ((e1 = op(e0, e4)) | ~ spl24_102), inference(avatar_component_clause, [], [f1117])).
fof(f18549, plain, (spl24_24 | ~ spl24_244 | ~ spl24_245), inference(avatar_split_clause, [], [f18548, f1852, f1848, f789])).
fof(f18548, plain, ((e3 = op(e4, e0)) | (~ spl24_244 | ~ spl24_245)), inference(forward_demodulation, [], [f1849, f1853])).
fof(f1849, plain, ((e3 = op(op(e0, op(e0, e0)), e0)) | ~ spl24_244), inference(avatar_component_clause, [], [f1848])).
fof(f18547, plain, (~ spl24_22 | ~ spl24_245 | spl24_259), inference(avatar_split_clause, [], [f18545, f1938, f1852, f781])).
fof(f1938, plain, (spl24_259 <=> (e1 = op(op(e0, op(e0, e0)), e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_259])])).
fof(f18545, plain, (~ (e1 = op(e4, e0)) | (~ spl24_245 | spl24_259)), inference(backward_demodulation, [], [f1940, f1853])).
fof(f1940, plain, (~ (e1 = op(op(e0, op(e0, e0)), e0)) | spl24_259), inference(avatar_component_clause, [], [f1938])).
fof(f18512, plain, (~ spl24_8 | ~ spl24_18), inference(avatar_split_clause, [], [f18506, f764, f722])).
fof(f18506, plain, (~ (e2 = op(e4, e3)) | ~ spl24_18), inference(backward_demodulation, [], [f477, f766])).
fof(f766, plain, ((e2 = op(e4, e1)) | ~ spl24_18), inference(avatar_component_clause, [], [f764])).
fof(f477, plain, ~ (op(e4, e1) = op(e4, e3)), inference(cnf_transformation, [], [f4])).
fof(f18496, plain, (~ spl24_31 | ~ spl24_41), inference(avatar_split_clause, [], [f18489, f861, f819])).
fof(f819, plain, (spl24_31 <=> (e0 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_31])])).
fof(f18489, plain, (~ (e0 = op(e3, e3)) | ~ spl24_41), inference(backward_demodulation, [], [f467, f863])).
fof(f863, plain, ((e0 = op(e3, e1)) | ~ spl24_41), inference(avatar_component_clause, [], [f861])).
fof(f467, plain, ~ (op(e3, e1) = op(e3, e3)), inference(cnf_transformation, [], [f4])).
fof(f18434, plain, (~ spl24_76 | ~ spl24_120 | spl24_181), inference(avatar_contradiction_clause, [], [f18433])).
fof(f18433, plain, ($false | (~ spl24_76 | ~ spl24_120 | spl24_181)), inference(subsumption_resolution, [], [f18430, f1194])).
fof(f1194, plain, ((e4 = op(e0, e1)) | ~ spl24_120), inference(avatar_component_clause, [], [f1192])).
fof(f1192, plain, (spl24_120 <=> (e4 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_120])])).
fof(f18430, plain, (~ (e4 = op(e0, e1)) | (~ spl24_76 | spl24_181)), inference(backward_demodulation, [], [f1533, f1010])).
fof(f1010, plain, ((e0 = op(e1, e4)) | ~ spl24_76), inference(avatar_component_clause, [], [f1008])).
fof(f1533, plain, (~ (e4 = op(op(e1, e4), e1)) | spl24_181), inference(avatar_component_clause, [], [f1531])).
fof(f1531, plain, (spl24_181 <=> (e4 = op(op(e1, e4), e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_181])])).
fof(f18432, plain, (~ spl24_51 | ~ spl24_76), inference(avatar_split_clause, [], [f18423, f1008, f903])).
fof(f18423, plain, (~ (e0 = op(e2, e4)) | ~ spl24_76), inference(backward_demodulation, [], [f426, f1010])).
fof(f426, plain, ~ (op(e1, e4) = op(e2, e4)), inference(cnf_transformation, [], [f4])).
fof(f18422, plain, (~ spl24_80 | ~ spl24_85), inference(avatar_split_clause, [], [f18417, f1045, f1024])).
fof(f18417, plain, (~ (e4 = op(e1, e4)) | ~ spl24_85), inference(backward_demodulation, [], [f451, f1047])).
fof(f1047, plain, ((e4 = op(e1, e3)) | ~ spl24_85), inference(avatar_component_clause, [], [f1045])).
fof(f18421, plain, (~ spl24_60 | ~ spl24_85), inference(avatar_split_clause, [], [f18414, f1045, f940])).
fof(f940, plain, (spl24_60 <=> (e4 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_60])])).
fof(f18414, plain, (~ (e4 = op(e2, e3)) | ~ spl24_85), inference(backward_demodulation, [], [f416, f1047])).
fof(f416, plain, ~ (op(e1, e3) = op(e2, e3)), inference(cnf_transformation, [], [f4])).
fof(f18412, plain, (~ spl24_37 | ~ spl24_87), inference(avatar_split_clause, [], [f18398, f1054, f844])).
fof(f844, plain, (spl24_37 <=> (e1 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_37])])).
fof(f1054, plain, (spl24_87 <=> (e1 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_87])])).
fof(f18398, plain, (~ (e1 = op(e3, e2)) | ~ spl24_87), inference(backward_demodulation, [], [f407, f1056])).
fof(f1056, plain, ((e1 = op(e1, e2)) | ~ spl24_87), inference(avatar_component_clause, [], [f1054])).
fof(f18394, plain, (~ spl24_84 | ~ spl24_94), inference(avatar_split_clause, [], [f18386, f1083, f1041])).
fof(f18386, plain, (~ (e3 = op(e1, e3)) | ~ spl24_94), inference(backward_demodulation, [], [f447, f1085])).
fof(f1085, plain, ((e3 = op(e1, e1)) | ~ spl24_94), inference(avatar_component_clause, [], [f1083])).
fof(f18381, plain, (~ spl24_72 | ~ spl24_98 | spl24_193), inference(avatar_split_clause, [], [f18379, f1592, f1100, f991])).
fof(f1592, plain, (spl24_193 <=> (e1 = op(op(e1, e0), e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_193])])).
fof(f18379, plain, (~ (e1 = op(e2, e0)) | (~ spl24_98 | spl24_193)), inference(backward_demodulation, [], [f1593, f1102])).
fof(f1102, plain, ((e2 = op(e1, e0)) | ~ spl24_98), inference(avatar_component_clause, [], [f1100])).
fof(f1593, plain, (~ (e1 = op(op(e1, e0), e0)) | spl24_193), inference(avatar_component_clause, [], [f1592])).
fof(f18380, plain, (~ spl24_88 | ~ spl24_98), inference(avatar_split_clause, [], [f18373, f1100, f1058])).
fof(f18373, plain, (~ (e2 = op(e1, e2)) | ~ spl24_98), inference(backward_demodulation, [], [f443, f1102])).
fof(f18370, plain, (spl24_76 | ~ spl24_102 | ~ spl24_197), inference(avatar_split_clause, [], [f18363, f1611, f1117, f1008])).
fof(f1611, plain, (spl24_197 <=> (e0 = op(op(e0, e4), e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_197])])).
fof(f18363, plain, ((e0 = op(e1, e4)) | (~ spl24_102 | ~ spl24_197)), inference(backward_demodulation, [], [f1613, f1119])).
fof(f1613, plain, ((e0 = op(op(e0, e4), e4)) | ~ spl24_197), inference(avatar_component_clause, [], [f1611])).
fof(f18356, plain, (~ spl24_8 | ~ spl24_108), inference(avatar_split_clause, [], [f18350, f1142, f722])).
fof(f18350, plain, (~ (e2 = op(e4, e3)) | ~ spl24_108), inference(backward_demodulation, [], [f415, f1144])).
fof(f1144, plain, ((e2 = op(e0, e3)) | ~ spl24_108), inference(avatar_component_clause, [], [f1142])).
fof(f18343, plain, (~ spl24_11 | ~ spl24_111), inference(avatar_split_clause, [], [f18328, f1155, f735])).
fof(f735, plain, (spl24_11 <=> (e0 = op(e4, e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_11])])).
fof(f1155, plain, (spl24_111 <=> (e0 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_111])])).
fof(f18328, plain, (~ (e0 = op(e4, e2)) | ~ spl24_111), inference(backward_demodulation, [], [f405, f1157])).
fof(f1157, plain, ((e0 = op(e0, e2)) | ~ spl24_111), inference(avatar_component_clause, [], [f1155])).
fof(f18324, plain, (spl24_22 | ~ spl24_120 | ~ spl24_205), inference(avatar_split_clause, [], [f18320, f1651, f1192, f781])).
fof(f1651, plain, (spl24_205 <=> (e1 = op(op(e0, e1), e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_205])])).
fof(f18320, plain, ((e1 = op(e4, e0)) | (~ spl24_120 | ~ spl24_205)), inference(backward_demodulation, [], [f1652, f1194])).
fof(f1652, plain, ((e1 = op(op(e0, e1), e0)) | ~ spl24_205), inference(avatar_component_clause, [], [f1651])).
fof(f18323, plain, (~ spl24_115 | ~ spl24_120), inference(avatar_split_clause, [], [f18316, f1192, f1171])).
fof(f18316, plain, (~ (e4 = op(e0, e2)) | ~ spl24_120), inference(backward_demodulation, [], [f436, f1194])).
fof(f436, plain, ~ (op(e0, e1) = op(e0, e2)), inference(cnf_transformation, [], [f4])).
fof(f18322, plain, (~ spl24_20 | ~ spl24_120), inference(avatar_split_clause, [], [f18315, f1192, f772])).
fof(f18315, plain, (~ (e4 = op(e4, e1)) | ~ spl24_120), inference(backward_demodulation, [], [f395, f1194])).
fof(f395, plain, ~ (op(e0, e1) = op(e4, e1)), inference(cnf_transformation, [], [f4])).
fof(f18311, plain, (~ spl24_107 | ~ spl24_124 | spl24_260), inference(avatar_split_clause, [], [f18307, f1944, f1209, f1138])).
fof(f1138, plain, (spl24_107 <=> (e1 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_107])])).
fof(f1944, plain, (spl24_260 <=> (e1 = op(e0, op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl24_260])])).
fof(f18307, plain, (~ (e1 = op(e0, e3)) | (~ spl24_124 | spl24_260)), inference(backward_demodulation, [], [f1946, f1211])).
fof(f1946, plain, (~ (e1 = op(e0, op(e0, e0))) | spl24_260), inference(avatar_component_clause, [], [f1944])).
fof(f18310, plain, (spl24_108 | ~ spl24_124 | ~ spl24_252), inference(avatar_split_clause, [], [f18306, f1888, f1209, f1142])).
fof(f1888, plain, (spl24_252 <=> (e2 = op(e0, op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl24_252])])).
fof(f18306, plain, ((e2 = op(e0, e3)) | (~ spl24_124 | ~ spl24_252)), inference(backward_demodulation, [], [f1889, f1211])).
fof(f1889, plain, ((e2 = op(e0, op(e0, e0))) | ~ spl24_252), inference(avatar_component_clause, [], [f1888])).
fof(f18309, plain, (~ spl24_104 | ~ spl24_124), inference(avatar_split_clause, [], [f18304, f1209, f1125])).
fof(f18304, plain, (~ (e3 = op(e0, e4)) | ~ spl24_124), inference(backward_demodulation, [], [f435, f1211])).
fof(f18297, plain, (spl24_15 | ~ spl24_128), inference(avatar_split_clause, [], [f18285, f1226, f751])).
fof(f1226, plain, (spl24_128 <=> (e2 = unit)), introduced(avatar_definition, [new_symbols(naming, [spl24_128])])).
fof(f18285, plain, ((e4 = op(e4, e2)) | ~ spl24_128), inference(backward_demodulation, [], [f330, f1228])).
fof(f1228, plain, ((e2 = unit) | ~ spl24_128), inference(avatar_component_clause, [], [f1226])).
fof(f18296, plain, (spl24_55 | ~ spl24_128), inference(avatar_split_clause, [], [f18284, f1226, f919])).
fof(f18284, plain, ((e4 = op(e2, e4)) | ~ spl24_128), inference(backward_demodulation, [], [f329, f1228])).
fof(f18295, plain, (spl24_39 | ~ spl24_128), inference(avatar_split_clause, [], [f18283, f1226, f852])).
fof(f852, plain, (spl24_39 <=> (e3 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_39])])).
fof(f18283, plain, ((e3 = op(e3, e2)) | ~ spl24_128), inference(backward_demodulation, [], [f328, f1228])).
fof(f18294, plain, (spl24_59 | ~ spl24_128), inference(avatar_split_clause, [], [f18282, f1226, f936])).
fof(f936, plain, (spl24_59 <=> (e3 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_59])])).
fof(f18282, plain, ((e3 = op(e2, e3)) | ~ spl24_128), inference(backward_demodulation, [], [f327, f1228])).
fof(f18293, plain, (spl24_63 | ~ spl24_128), inference(avatar_split_clause, [], [f18281, f1226, f953])).
fof(f953, plain, (spl24_63 <=> (e2 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_63])])).
fof(f18281, plain, ((e2 = op(e2, e2)) | ~ spl24_128), inference(backward_demodulation, [], [f326, f1228])).
fof(f18291, plain, (spl24_87 | ~ spl24_128), inference(avatar_split_clause, [], [f18279, f1226, f1054])).
fof(f18279, plain, ((e1 = op(e1, e2)) | ~ spl24_128), inference(backward_demodulation, [], [f324, f1228])).
fof(f18290, plain, (spl24_67 | ~ spl24_128), inference(avatar_split_clause, [], [f18278, f1226, f970])).
fof(f18278, plain, ((e1 = op(e2, e1)) | ~ spl24_128), inference(backward_demodulation, [], [f323, f1228])).
fof(f18289, plain, (spl24_111 | ~ spl24_128), inference(avatar_split_clause, [], [f18277, f1226, f1155])).
fof(f18277, plain, ((e0 = op(e0, e2)) | ~ spl24_128), inference(backward_demodulation, [], [f322, f1228])).
fof(f18288, plain, (spl24_71 | ~ spl24_128), inference(avatar_split_clause, [], [f18276, f1226, f987])).
fof(f987, plain, (spl24_71 <=> (e0 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_71])])).
fof(f18276, plain, ((e0 = op(e2, e0)) | ~ spl24_128), inference(backward_demodulation, [], [f321, f1228])).
fof(f18275, plain, (~ spl24_8 | ~ spl24_28 | spl24_134), inference(avatar_split_clause, [], [f18274, f1303, f806, f722])).
fof(f806, plain, (spl24_28 <=> (e2 = op(e3, e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_28])])).
fof(f18274, plain, (~ (e2 = op(e4, e3)) | (~ spl24_28 | spl24_134)), inference(forward_demodulation, [], [f1305, f808])).
fof(f808, plain, ((e2 = op(e3, e4)) | ~ spl24_28), inference(avatar_component_clause, [], [f806])).
fof(f1305, plain, (~ (op(e3, e4) = op(e4, e3)) | spl24_134), inference(avatar_component_clause, [], [f1303])).
fof(f18270, plain, (~ spl24_60 | ~ spl24_28 | spl24_148), inference(avatar_split_clause, [], [f18269, f1369, f806, f940])).
fof(f1369, plain, (spl24_148 <=> (e4 = op(op(e3, e4), e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_148])])).
fof(f18269, plain, (~ (e4 = op(e2, e3)) | (~ spl24_28 | spl24_148)), inference(forward_demodulation, [], [f1371, f808])).
fof(f1371, plain, (~ (e4 = op(op(e3, e4), e3)) | spl24_148), inference(avatar_component_clause, [], [f1369])).
fof(f18222, plain, (~ spl24_72 | ~ spl24_252 | spl24_259), inference(avatar_split_clause, [], [f18220, f1938, f1888, f991])).
fof(f18220, plain, (~ (e1 = op(e2, e0)) | (~ spl24_252 | spl24_259)), inference(backward_demodulation, [], [f1940, f1889])).
fof(f18201, plain, (~ spl24_75 | ~ spl24_50), inference(avatar_split_clause, [], [f17990, f898, f1003])).
fof(f17990, plain, (~ (e4 = op(e2, e0)) | ~ spl24_50), inference(forward_demodulation, [], [f389, f900])).
fof(f900, plain, ((e4 = op(e3, e0)) | ~ spl24_50), inference(avatar_component_clause, [], [f898])).
fof(f389, plain, ~ (op(e2, e0) = op(e3, e0)), inference(cnf_transformation, [], [f4])).
fof(f18198, plain, (~ spl24_45 | ~ spl24_50), inference(avatar_split_clause, [], [f18197, f898, f877])).
fof(f877, plain, (spl24_45 <=> (e4 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_45])])).
fof(f18197, plain, (~ (e4 = op(e3, e1)) | ~ spl24_50), inference(forward_demodulation, [], [f462, f900])).
fof(f462, plain, ~ (op(e3, e0) = op(e3, e1)), inference(cnf_transformation, [], [f4])).
fof(f18195, plain, (~ spl24_40 | ~ spl24_50), inference(avatar_split_clause, [], [f18194, f898, f856])).
fof(f18194, plain, (~ (e4 = op(e3, e2)) | ~ spl24_50), inference(forward_demodulation, [], [f463, f900])).
fof(f18164, plain, (~ spl24_28 | ~ spl24_37 | ~ spl24_134 | ~ spl24_212 | spl24_238), inference(avatar_contradiction_clause, [], [f18163])).
fof(f18163, plain, ($false | (~ spl24_28 | ~ spl24_37 | ~ spl24_134 | ~ spl24_212 | spl24_238)), inference(subsumption_resolution, [], [f18158, f846])).
fof(f846, plain, ((e1 = op(e3, e2)) | ~ spl24_37), inference(avatar_component_clause, [], [f844])).
fof(f18158, plain, (~ (e1 = op(e3, e2)) | (~ spl24_28 | ~ spl24_134 | ~ spl24_212 | spl24_238)), inference(backward_demodulation, [], [f18082, f808])).
fof(f18082, plain, (~ (e1 = op(e3, op(e3, e4))) | (~ spl24_134 | ~ spl24_212 | spl24_238)), inference(backward_demodulation, [], [f18061, f1304])).
fof(f18061, plain, (~ (e1 = op(e3, op(e4, e3))) | (~ spl24_212 | spl24_238)), inference(backward_demodulation, [], [f1818, f1689])).
fof(f1689, plain, ((e4 = op(e3, op(e3, e3))) | ~ spl24_212), inference(avatar_component_clause, [], [f1688])).
fof(f1688, plain, (spl24_212 <=> (e4 = op(e3, op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl24_212])])).
fof(f18161, plain, (spl24_8 | ~ spl24_28 | ~ spl24_134), inference(avatar_split_clause, [], [f18154, f1303, f806, f722])).
fof(f18154, plain, ((e2 = op(e4, e3)) | (~ spl24_28 | ~ spl24_134)), inference(backward_demodulation, [], [f1304, f808])).
fof(f18160, plain, (~ spl24_3 | ~ spl24_28), inference(avatar_split_clause, [], [f18151, f806, f701])).
fof(f18151, plain, (~ (e2 = op(e4, e4)) | ~ spl24_28), inference(backward_demodulation, [], [f431, f808])).
fof(f18150, plain, (~ spl24_26 | ~ spl24_31), inference(avatar_split_clause, [], [f18141, f819, f798])).
fof(f18141, plain, (~ (e0 = op(e3, e4)) | ~ spl24_31), inference(backward_demodulation, [], [f471, f821])).
fof(f821, plain, ((e0 = op(e3, e3)) | ~ spl24_31), inference(avatar_component_clause, [], [f819])).
fof(f471, plain, ~ (op(e3, e3) = op(e3, e4)), inference(cnf_transformation, [], [f4])).
fof(f18130, plain, (spl24_8 | ~ spl24_60 | ~ spl24_169), inference(avatar_split_clause, [], [f18128, f1472, f940, f722])).
fof(f1472, plain, (spl24_169 <=> (e2 = op(op(e2, e3), e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_169])])).
fof(f18128, plain, ((e2 = op(e4, e3)) | (~ spl24_60 | ~ spl24_169)), inference(backward_demodulation, [], [f1474, f942])).
fof(f942, plain, ((e4 = op(e2, e3)) | ~ spl24_60), inference(avatar_component_clause, [], [f940])).
fof(f1474, plain, ((e2 = op(op(e2, e3), e3)) | ~ spl24_169), inference(avatar_component_clause, [], [f1472])).
fof(f18123, plain, (~ spl24_56 | ~ spl24_64 | spl24_253), inference(avatar_split_clause, [], [f18119, f1893, f957, f924])).
fof(f924, plain, (spl24_56 <=> (e0 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_56])])).
fof(f1893, plain, (spl24_253 <=> (e0 = op(e2, op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl24_253])])).
fof(f18119, plain, (~ (e0 = op(e2, e3)) | (~ spl24_64 | spl24_253)), inference(backward_demodulation, [], [f1895, f959])).
fof(f959, plain, ((e3 = op(e2, e2)) | ~ spl24_64), inference(avatar_component_clause, [], [f957])).
fof(f1895, plain, (~ (e0 = op(e2, op(e2, e2))) | spl24_253), inference(avatar_component_clause, [], [f1893])).
fof(f18120, plain, (~ spl24_54 | ~ spl24_64), inference(avatar_split_clause, [], [f18115, f957, f915])).
fof(f915, plain, (spl24_54 <=> (e3 = op(e2, e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_54])])).
fof(f18115, plain, (~ (e3 = op(e2, e4)) | ~ spl24_64), inference(backward_demodulation, [], [f460, f959])).
fof(f18113, plain, (~ spl24_26 | ~ spl24_104 | spl24_197), inference(avatar_split_clause, [], [f18111, f1611, f1125, f798])).
fof(f18111, plain, (~ (e0 = op(e3, e4)) | (~ spl24_104 | spl24_197)), inference(backward_demodulation, [], [f1612, f1127])).
fof(f1127, plain, ((e3 = op(e0, e4)) | ~ spl24_104), inference(avatar_component_clause, [], [f1125])).
fof(f1612, plain, (~ (e0 = op(op(e0, e4), e4)) | spl24_197), inference(avatar_component_clause, [], [f1611])).
fof(f18112, plain, (~ spl24_54 | ~ spl24_104), inference(avatar_split_clause, [], [f18106, f1125, f915])).
fof(f18106, plain, (~ (e3 = op(e2, e4)) | ~ spl24_104), inference(backward_demodulation, [], [f423, f1127])).
fof(f423, plain, ~ (op(e0, e4) = op(e2, e4)), inference(cnf_transformation, [], [f4])).
fof(f18104, plain, (~ spl24_102 | ~ spl24_107), inference(avatar_split_clause, [], [f18097, f1138, f1117])).
fof(f18097, plain, (~ (e1 = op(e0, e4)) | ~ spl24_107), inference(backward_demodulation, [], [f441, f1140])).
fof(f1140, plain, ((e1 = op(e0, e3)) | ~ spl24_107), inference(avatar_component_clause, [], [f1138])).
fof(f441, plain, ~ (op(e0, e3) = op(e0, e4)), inference(cnf_transformation, [], [f4])).
fof(f18103, plain, (~ spl24_7 | ~ spl24_107), inference(avatar_split_clause, [], [f18096, f1138, f718])).
fof(f718, plain, (spl24_7 <=> (e1 = op(e4, e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_7])])).
fof(f18096, plain, (~ (e1 = op(e4, e3)) | ~ spl24_107), inference(backward_demodulation, [], [f415, f1140])).
fof(f18093, plain, (~ spl24_110 | ~ spl24_115), inference(avatar_split_clause, [], [f18086, f1171, f1150])).
fof(f18086, plain, (~ (e4 = op(e0, e3)) | ~ spl24_115), inference(backward_demodulation, [], [f439, f1173])).
fof(f439, plain, ~ (op(e0, e2) = op(e0, e3)), inference(cnf_transformation, [], [f4])).
fof(f18092, plain, (~ spl24_65 | ~ spl24_115), inference(avatar_split_clause, [], [f18085, f1171, f961])).
fof(f961, plain, (spl24_65 <=> (e4 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_65])])).
fof(f18085, plain, (~ (e4 = op(e2, e2)) | ~ spl24_115), inference(backward_demodulation, [], [f403, f1173])).
fof(f18078, plain, (spl24_51 | ~ spl24_11 | ~ spl24_138), inference(avatar_split_clause, [], [f18077, f1322, f735, f903])).
fof(f18077, plain, ((e0 = op(e2, e4)) | (~ spl24_11 | ~ spl24_138)), inference(forward_demodulation, [], [f1323, f737])).
fof(f737, plain, ((e0 = op(e4, e2)) | ~ spl24_11), inference(avatar_component_clause, [], [f735])).
fof(f1323, plain, ((op(e2, e4) = op(e4, e2)) | ~ spl24_138), inference(avatar_component_clause, [], [f1322])).
fof(f18074, plain, (~ spl24_26 | ~ spl24_24 | spl24_144), inference(avatar_split_clause, [], [f18073, f1350, f789, f798])).
fof(f1350, plain, (spl24_144 <=> (e0 = op(op(e4, e0), e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_144])])).
fof(f18073, plain, (~ (e0 = op(e3, e4)) | (~ spl24_24 | spl24_144)), inference(forward_demodulation, [], [f1352, f791])).
fof(f1352, plain, (~ (e0 = op(op(e4, e0), e4)) | spl24_144), inference(avatar_component_clause, [], [f1350])).
fof(f18072, plain, (spl24_104 | ~ spl24_24 | ~ spl24_146), inference(avatar_split_clause, [], [f18071, f1360, f789, f1125])).
fof(f18071, plain, ((e3 = op(e0, e4)) | (~ spl24_24 | ~ spl24_146)), inference(forward_demodulation, [], [f1361, f791])).
fof(f1361, plain, ((op(e0, e4) = op(e4, e0)) | ~ spl24_146), inference(avatar_component_clause, [], [f1360])).
fof(f18070, plain, (~ spl24_110 | ~ spl24_50 | spl24_163), inference(avatar_split_clause, [], [f18069, f1443, f898, f1150])).
fof(f1443, plain, (spl24_163 <=> (op(e0, e3) = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_163])])).
fof(f18069, plain, (~ (e4 = op(e0, e3)) | (~ spl24_50 | spl24_163)), inference(forward_demodulation, [], [f1445, f900])).
fof(f1445, plain, (~ (op(e0, e3) = op(e3, e0)) | spl24_163), inference(avatar_component_clause, [], [f1443])).
fof(f18064, plain, (~ spl24_7 | ~ spl24_212 | spl24_225), inference(avatar_split_clause, [], [f18060, f1746, f1688, f718])).
fof(f1746, plain, (spl24_225 <=> (e1 = op(op(e3, op(e3, e3)), e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_225])])).
fof(f18060, plain, (~ (e1 = op(e4, e3)) | (~ spl24_212 | spl24_225)), inference(backward_demodulation, [], [f1748, f1689])).
fof(f1748, plain, (~ (e1 = op(op(e3, op(e3, e3)), e3)) | spl24_225), inference(avatar_component_clause, [], [f1746])).
fof(f18039, plain, (~ spl24_114 | ~ spl24_123 | spl24_248), inference(avatar_split_clause, [], [f18038, f1866, f1205, f1167])).
fof(f18038, plain, (~ (e3 = op(e0, e2)) | (~ spl24_123 | spl24_248)), inference(forward_demodulation, [], [f1868, f1207])).
fof(f17989, plain, (spl24_115 | ~ spl24_11 | ~ spl24_137), inference(avatar_split_clause, [], [f17724, f1317, f735, f1171])).
fof(f17724, plain, ((e4 = op(e0, e2)) | (~ spl24_11 | ~ spl24_137)), inference(forward_demodulation, [], [f1319, f737])).
fof(f17982, plain, (~ spl24_57 | ~ spl24_37 | spl24_155), inference(avatar_split_clause, [], [f17981, f1405, f844, f928])).
fof(f1405, plain, (spl24_155 <=> (op(e2, e3) = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_155])])).
fof(f17981, plain, (~ (e1 = op(e2, e3)) | (~ spl24_37 | spl24_155)), inference(forward_demodulation, [], [f1407, f846])).
fof(f1407, plain, (~ (op(e2, e3) = op(e3, e2)) | spl24_155), inference(avatar_component_clause, [], [f1405])).
fof(f17980, plain, (~ spl24_58 | ~ spl24_68), inference(avatar_split_clause, [], [f17979, f974, f932])).
fof(f17979, plain, (~ (e2 = op(e2, e3)) | ~ spl24_68), inference(forward_demodulation, [], [f457, f976])).
fof(f976, plain, ((e2 = op(e2, e1)) | ~ spl24_68), inference(avatar_component_clause, [], [f974])).
fof(f457, plain, ~ (op(e2, e1) = op(e2, e3)), inference(cnf_transformation, [], [f4])).
fof(f17978, plain, (~ spl24_59 | ~ spl24_84), inference(avatar_split_clause, [], [f17977, f1041, f936])).
fof(f17977, plain, (~ (e3 = op(e2, e3)) | ~ spl24_84), inference(forward_demodulation, [], [f416, f1043])).
fof(f17976, plain, (~ spl24_57 | ~ spl24_72), inference(avatar_split_clause, [], [f17975, f991, f928])).
fof(f17975, plain, (~ (e1 = op(e2, e3)) | ~ spl24_72), inference(forward_demodulation, [], [f454, f993])).
fof(f454, plain, ~ (op(e2, e0) = op(e2, e3)), inference(cnf_transformation, [], [f4])).
fof(f17973, plain, (~ spl24_55 | ~ spl24_80), inference(avatar_split_clause, [], [f17972, f1024, f919])).
fof(f17972, plain, (~ (e4 = op(e2, e4)) | ~ spl24_80), inference(forward_demodulation, [], [f426, f1026])).
fof(f17969, plain, (~ spl24_34 | ~ spl24_44), inference(avatar_split_clause, [], [f17968, f873, f831])).
fof(f17968, plain, (~ (e3 = op(e3, e3)) | ~ spl24_44), inference(forward_demodulation, [], [f467, f875])).
fof(f875, plain, ((e3 = op(e3, e1)) | ~ spl24_44), inference(avatar_component_clause, [], [f873])).
fof(f17965, plain, (~ spl24_10 | ~ spl24_20), inference(avatar_split_clause, [], [f17964, f772, f730])).
fof(f17964, plain, (~ (e4 = op(e4, e3)) | ~ spl24_20), inference(forward_demodulation, [], [f477, f774])).
fof(f17963, plain, (~ spl24_9 | ~ spl24_24), inference(avatar_split_clause, [], [f17962, f789, f726])).
fof(f726, plain, (spl24_9 <=> (e3 = op(e4, e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_9])])).
fof(f17962, plain, (~ (e3 = op(e4, e3)) | ~ spl24_24), inference(forward_demodulation, [], [f474, f791])).
fof(f474, plain, ~ (op(e4, e0) = op(e4, e3)), inference(cnf_transformation, [], [f4])).
fof(f17961, plain, (~ spl24_9 | ~ spl24_84), inference(avatar_split_clause, [], [f17960, f1041, f726])).
fof(f17960, plain, (~ (e3 = op(e4, e3)) | ~ spl24_84), inference(forward_demodulation, [], [f418, f1043])).
fof(f418, plain, ~ (op(e1, e3) = op(e4, e3)), inference(cnf_transformation, [], [f4])).
fof(f17958, plain, (~ spl24_5 | ~ spl24_20), inference(avatar_split_clause, [], [f17957, f772, f709])).
fof(f709, plain, (spl24_5 <=> (e4 = op(e4, e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_5])])).
fof(f17957, plain, (~ (e4 = op(e4, e4)) | ~ spl24_20), inference(forward_demodulation, [], [f478, f774])).
fof(f478, plain, ~ (op(e4, e1) = op(e4, e4)), inference(cnf_transformation, [], [f4])).
fof(f17956, plain, (~ spl24_5 | ~ spl24_80), inference(avatar_split_clause, [], [f17955, f1024, f709])).
fof(f17955, plain, (~ (e4 = op(e4, e4)) | ~ spl24_80), inference(forward_demodulation, [], [f428, f1026])).
fof(f17930, plain, (~ spl24_27 | ~ spl24_37), inference(avatar_split_clause, [], [f17918, f844, f802])).
fof(f17918, plain, (~ (e1 = op(e3, e4)) | ~ spl24_37), inference(backward_demodulation, [], [f470, f846])).
fof(f17916, plain, (~ spl24_40 | ~ spl24_65), inference(avatar_split_clause, [], [f17912, f961, f856])).
fof(f17912, plain, (~ (e4 = op(e3, e2)) | ~ spl24_65), inference(backward_demodulation, [], [f409, f963])).
fof(f963, plain, ((e4 = op(e2, e2)) | ~ spl24_65), inference(avatar_component_clause, [], [f961])).
fof(f409, plain, ~ (op(e2, e2) = op(e3, e2)), inference(cnf_transformation, [], [f4])).
fof(f17908, plain, (~ spl24_62 | ~ spl24_72), inference(avatar_split_clause, [], [f17902, f991, f949])).
fof(f949, plain, (spl24_62 <=> (e1 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_62])])).
fof(f17902, plain, (~ (e1 = op(e2, e2)) | ~ spl24_72), inference(backward_demodulation, [], [f453, f993])).
fof(f453, plain, ~ (op(e2, e0) = op(e2, e2)), inference(cnf_transformation, [], [f4])).
fof(f17858, plain, (~ spl24_92 | ~ spl24_232), inference(avatar_contradiction_clause, [], [f17857])).
fof(f17857, plain, ($false | (~ spl24_92 | ~ spl24_232)), inference(subsumption_resolution, [], [f17856, f487])).
fof(f487, plain, ~ (e1 = e3), inference(cnf_transformation, [], [f5])).
fof(f17856, plain, ((e1 = e3) | (~ spl24_92 | ~ spl24_232)), inference(forward_demodulation, [], [f17853, f1077])).
fof(f17853, plain, ((e3 = op(e1, e1)) | (~ spl24_92 | ~ spl24_232)), inference(backward_demodulation, [], [f1779, f1077])).
fof(f1779, plain, ((e3 = op(e1, op(e1, e1))) | ~ spl24_232), inference(avatar_component_clause, [], [f1778])).
fof(f1778, plain, (spl24_232 <=> (e3 = op(e1, op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl24_232])])).
fof(f17839, plain, (~ spl24_19 | ~ spl24_3 | ~ spl24_11 | ~ spl24_102 | spl24_263), inference(avatar_split_clause, [], [f17835, f1965, f1117, f735, f701, f768])).
fof(f768, plain, (spl24_19 <=> (e3 = op(e4, e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_19])])).
fof(f1965, plain, (spl24_263 <=> (e3 = op(e4, op(op(e4, op(e4, e4)), e4)))), introduced(avatar_definition, [new_symbols(naming, [spl24_263])])).
fof(f17835, plain, (~ (e3 = op(e4, e1)) | (~ spl24_3 | ~ spl24_11 | ~ spl24_102 | spl24_263)), inference(backward_demodulation, [], [f17649, f1119])).
fof(f17649, plain, (~ (e3 = op(e4, op(e0, e4))) | (~ spl24_3 | ~ spl24_11 | spl24_263)), inference(forward_demodulation, [], [f17648, f737])).
fof(f17648, plain, (~ (e3 = op(e4, op(op(e4, e2), e4))) | (~ spl24_3 | spl24_263)), inference(forward_demodulation, [], [f1967, f703])).
fof(f703, plain, ((e2 = op(e4, e4)) | ~ spl24_3), inference(avatar_component_clause, [], [f701])).
fof(f1967, plain, (~ (e3 = op(e4, op(op(e4, op(e4, e4)), e4))) | spl24_263), inference(avatar_component_clause, [], [f1965])).
fof(f17837, plain, (~ spl24_27 | ~ spl24_102), inference(avatar_split_clause, [], [f17828, f1117, f802])).
fof(f17828, plain, (~ (e1 = op(e3, e4)) | ~ spl24_102), inference(backward_demodulation, [], [f424, f1119])).
fof(f424, plain, ~ (op(e0, e4) = op(e3, e4)), inference(cnf_transformation, [], [f4])).
fof(f17820, plain, (~ spl24_105 | ~ spl24_110), inference(avatar_split_clause, [], [f17806, f1150, f1129])).
fof(f1129, plain, (spl24_105 <=> (e4 = op(e0, e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_105])])).
fof(f17806, plain, (~ (e4 = op(e0, e4)) | ~ spl24_110), inference(backward_demodulation, [], [f441, f1152])).
fof(f17819, plain, (~ spl24_85 | ~ spl24_110), inference(avatar_split_clause, [], [f17805, f1150, f1045])).
fof(f17805, plain, (~ (e4 = op(e1, e3)) | ~ spl24_110), inference(backward_demodulation, [], [f412, f1152])).
fof(f412, plain, ~ (op(e0, e3) = op(e1, e3)), inference(cnf_transformation, [], [f4])).
fof(f17804, plain, (~ spl24_109 | ~ spl24_114), inference(avatar_split_clause, [], [f17797, f1167, f1146])).
fof(f1146, plain, (spl24_109 <=> (e3 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_109])])).
fof(f17797, plain, (~ (e3 = op(e0, e3)) | ~ spl24_114), inference(backward_demodulation, [], [f439, f1169])).
fof(f17803, plain, (~ spl24_89 | ~ spl24_114), inference(avatar_split_clause, [], [f17794, f1167, f1062])).
fof(f1062, plain, (spl24_89 <=> (e3 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_89])])).
fof(f17794, plain, (~ (e3 = op(e1, e2)) | ~ spl24_114), inference(backward_demodulation, [], [f402, f1169])).
fof(f402, plain, ~ (op(e0, e2) = op(e1, e2)), inference(cnf_transformation, [], [f4])).
fof(f17793, plain, (~ spl24_76 | ~ spl24_116 | ~ spl24_194), inference(avatar_split_clause, [], [f17788, f1597, f1176, f1008])).
fof(f17788, plain, (~ (e0 = op(e1, e4)) | (~ spl24_116 | ~ spl24_194)), inference(backward_demodulation, [], [f16957, f1178])).
fof(f16957, plain, (~ (op(e0, e1) = op(e1, e4)) | ~ spl24_194), inference(backward_demodulation, [], [f445, f1598])).
fof(f445, plain, ~ (op(e1, e0) = op(e1, e4)), inference(cnf_transformation, [], [f4])).
fof(f17769, plain, (spl24_114 | ~ spl24_123 | ~ spl24_248), inference(avatar_split_clause, [], [f17764, f1866, f1205, f1167])).
fof(f17764, plain, ((e3 = op(e0, e2)) | (~ spl24_123 | ~ spl24_248)), inference(backward_demodulation, [], [f1867, f1207])).
fof(f1867, plain, ((e3 = op(e0, op(e0, e0))) | ~ spl24_248), inference(avatar_component_clause, [], [f1866])).
fof(f17767, plain, (~ spl24_73 | ~ spl24_123), inference(avatar_split_clause, [], [f17755, f1205, f995])).
fof(f995, plain, (spl24_73 <=> (e2 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_73])])).
fof(f17755, plain, (~ (e2 = op(e2, e0)) | ~ spl24_123), inference(backward_demodulation, [], [f383, f1207])).
fof(f17728, plain, (~ spl24_27 | ~ spl24_7 | spl24_134), inference(avatar_split_clause, [], [f17727, f1303, f718, f802])).
fof(f17727, plain, (~ (e1 = op(e3, e4)) | (~ spl24_7 | spl24_134)), inference(forward_demodulation, [], [f1305, f720])).
fof(f720, plain, ((e1 = op(e4, e3)) | ~ spl24_7), inference(avatar_component_clause, [], [f718])).
fof(f17712, plain, (~ spl24_40 | ~ spl24_54 | spl24_165), inference(avatar_split_clause, [], [f17711, f1452, f915, f856])).
fof(f1452, plain, (spl24_165 <=> (e4 = op(op(e2, e4), e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_165])])).
fof(f17711, plain, (~ (e4 = op(e3, e2)) | (~ spl24_54 | spl24_165)), inference(forward_demodulation, [], [f1454, f917])).
fof(f917, plain, ((e3 = op(e2, e4)) | ~ spl24_54), inference(avatar_component_clause, [], [f915])).
fof(f1454, plain, (~ (e4 = op(op(e2, e4), e2)) | spl24_165), inference(avatar_component_clause, [], [f1452])).
fof(f17710, plain, (spl24_114 | ~ spl24_56 | ~ spl24_168), inference(avatar_split_clause, [], [f17709, f1467, f924, f1167])).
fof(f1467, plain, (spl24_168 <=> (e3 = op(op(e2, e3), e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_168])])).
fof(f17709, plain, ((e3 = op(e0, e2)) | (~ spl24_56 | ~ spl24_168)), inference(forward_demodulation, [], [f1468, f926])).
fof(f926, plain, ((e0 = op(e2, e3)) | ~ spl24_56), inference(avatar_component_clause, [], [f924])).
fof(f1468, plain, ((e3 = op(op(e2, e3), e2)) | ~ spl24_168), inference(avatar_component_clause, [], [f1467])).
fof(f17700, plain, (~ spl24_40 | ~ spl24_33 | spl24_212), inference(avatar_split_clause, [], [f17699, f1688, f827, f856])).
fof(f17699, plain, (~ (e4 = op(e3, e2)) | (~ spl24_33 | spl24_212)), inference(forward_demodulation, [], [f1690, f829])).
fof(f1690, plain, (~ (e4 = op(e3, op(e3, e3))) | spl24_212), inference(avatar_component_clause, [], [f1688])).
fof(f17614, plain, (~ spl24_3 | ~ spl24_11 | spl24_246), inference(avatar_contradiction_clause, [], [f17613])).
fof(f17613, plain, ($false | (~ spl24_3 | ~ spl24_11 | spl24_246)), inference(subsumption_resolution, [], [f17605, f737])).
fof(f17605, plain, (~ (e0 = op(e4, e2)) | (~ spl24_3 | spl24_246)), inference(backward_demodulation, [], [f1859, f703])).
fof(f1859, plain, (~ (e0 = op(e4, op(e4, e4))) | spl24_246), inference(avatar_component_clause, [], [f1857])).
fof(f1857, plain, (spl24_246 <=> (e0 = op(e4, op(e4, e4)))), introduced(avatar_definition, [new_symbols(naming, [spl24_246])])).
fof(f17596, plain, (~ spl24_1 | ~ spl24_11), inference(avatar_split_clause, [], [f17589, f735, f693])).
fof(f17589, plain, (~ (e0 = op(e4, e4)) | ~ spl24_11), inference(backward_demodulation, [], [f480, f737])).
fof(f480, plain, ~ (op(e4, e2) = op(e4, e4)), inference(cnf_transformation, [], [f4])).
fof(f17587, plain, (~ spl24_14 | ~ spl24_19), inference(avatar_split_clause, [], [f17582, f768, f747])).
fof(f17582, plain, (~ (e3 = op(e4, e2)) | ~ spl24_19), inference(backward_demodulation, [], [f476, f770])).
fof(f770, plain, ((e3 = op(e4, e1)) | ~ spl24_19), inference(avatar_component_clause, [], [f768])).
fof(f17574, plain, (~ spl24_28 | ~ spl24_33), inference(avatar_split_clause, [], [f17567, f827, f806])).
fof(f17567, plain, (~ (e2 = op(e3, e4)) | ~ spl24_33), inference(backward_demodulation, [], [f471, f829])).
fof(f17545, plain, (~ spl24_14 | ~ spl24_54 | spl24_138), inference(avatar_split_clause, [], [f17542, f1322, f915, f747])).
fof(f17542, plain, (~ (e3 = op(e4, e2)) | (~ spl24_54 | spl24_138)), inference(backward_demodulation, [], [f1324, f917])).
fof(f17537, plain, (~ spl24_31 | ~ spl24_56), inference(avatar_split_clause, [], [f17528, f924, f819])).
fof(f17528, plain, (~ (e0 = op(e3, e3)) | ~ spl24_56), inference(backward_demodulation, [], [f419, f926])).
fof(f17527, plain, (~ spl24_52 | ~ spl24_62), inference(avatar_split_clause, [], [f17518, f949, f907])).
fof(f17518, plain, (~ (e1 = op(e2, e4)) | ~ spl24_62), inference(backward_demodulation, [], [f460, f951])).
fof(f951, plain, ((e1 = op(e2, e2)) | ~ spl24_62), inference(avatar_component_clause, [], [f949])).
fof(f17526, plain, (~ spl24_37 | ~ spl24_62), inference(avatar_split_clause, [], [f17515, f949, f844])).
fof(f17515, plain, (~ (e1 = op(e3, e2)) | ~ spl24_62), inference(backward_demodulation, [], [f409, f951])).
fof(f17513, plain, (~ spl24_60 | ~ spl24_70), inference(avatar_split_clause, [], [f17508, f982, f940])).
fof(f17508, plain, (~ (e4 = op(e2, e3)) | ~ spl24_70), inference(backward_demodulation, [], [f457, f984])).
fof(f17512, plain, (~ spl24_45 | ~ spl24_70), inference(avatar_split_clause, [], [f17505, f982, f877])).
fof(f17505, plain, (~ (e4 = op(e3, e1)) | ~ spl24_70), inference(backward_demodulation, [], [f399, f984])).
fof(f399, plain, ~ (op(e2, e1) = op(e3, e1)), inference(cnf_transformation, [], [f4])).
fof(f17504, plain, (~ spl24_1 | ~ spl24_76), inference(avatar_split_clause, [], [f17497, f1008, f693])).
fof(f17497, plain, (~ (e0 = op(e4, e4)) | ~ spl24_76), inference(backward_demodulation, [], [f428, f1010])).
fof(f17493, plain, (~ spl24_45 | ~ spl24_85 | spl24_159), inference(avatar_split_clause, [], [f17489, f1424, f1045, f877])).
fof(f1424, plain, (spl24_159 <=> (op(e1, e3) = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_159])])).
fof(f17489, plain, (~ (e4 = op(e3, e1)) | (~ spl24_85 | spl24_159)), inference(backward_demodulation, [], [f1426, f1047])).
fof(f1426, plain, (~ (op(e1, e3) = op(e3, e1)) | spl24_159), inference(avatar_component_clause, [], [f1424])).
fof(f17484, plain, (~ spl24_69 | ~ spl24_89 | spl24_175), inference(avatar_split_clause, [], [f17479, f1503, f1062, f978])).
fof(f978, plain, (spl24_69 <=> (e3 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_69])])).
fof(f1503, plain, (spl24_175 <=> (op(e1, e2) = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_175])])).
fof(f17479, plain, (~ (e3 = op(e2, e1)) | (~ spl24_89 | spl24_175)), inference(backward_demodulation, [], [f1505, f1064])).
fof(f1064, plain, ((e3 = op(e1, e2)) | ~ spl24_89), inference(avatar_component_clause, [], [f1062])).
fof(f1505, plain, (~ (op(e1, e2) = op(e2, e1)) | spl24_175), inference(avatar_component_clause, [], [f1503])).
fof(f17483, plain, (~ spl24_79 | ~ spl24_89), inference(avatar_split_clause, [], [f17477, f1062, f1020])).
fof(f17477, plain, (~ (e3 = op(e1, e4)) | ~ spl24_89), inference(backward_demodulation, [], [f450, f1064])).
fof(f17482, plain, (~ spl24_14 | ~ spl24_89), inference(avatar_split_clause, [], [f17475, f1062, f747])).
fof(f17475, plain, (~ (e3 = op(e4, e2)) | ~ spl24_89), inference(backward_demodulation, [], [f408, f1064])).
fof(f17472, plain, (spl24_89 | ~ spl24_93 | ~ spl24_232), inference(avatar_split_clause, [], [f17467, f1778, f1079, f1062])).
fof(f17467, plain, ((e3 = op(e1, e2)) | (~ spl24_93 | ~ spl24_232)), inference(backward_demodulation, [], [f1779, f1081])).
fof(f17471, plain, (~ spl24_83 | ~ spl24_93), inference(avatar_split_clause, [], [f17463, f1079, f1037])).
fof(f1037, plain, (spl24_83 <=> (e2 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_83])])).
fof(f17463, plain, (~ (e2 = op(e1, e3)) | ~ spl24_93), inference(backward_demodulation, [], [f447, f1081])).
fof(f17458, plain, (~ spl24_79 | ~ spl24_7 | spl24_132), inference(avatar_split_clause, [], [f17457, f1293, f718, f1020])).
fof(f1293, plain, (spl24_132 <=> (e3 = op(op(e4, e3), e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_132])])).
fof(f17457, plain, (~ (e3 = op(e1, e4)) | (~ spl24_7 | spl24_132)), inference(forward_demodulation, [], [f1295, f720])).
fof(f1295, plain, (~ (e3 = op(op(e4, e3), e4)) | spl24_132), inference(avatar_component_clause, [], [f1293])).
fof(f17456, plain, (spl24_85 | ~ spl24_7 | ~ spl24_133), inference(avatar_split_clause, [], [f17455, f1298, f718, f1045])).
fof(f1298, plain, (spl24_133 <=> (e4 = op(op(e4, e3), e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_133])])).
fof(f17455, plain, ((e4 = op(e1, e3)) | (~ spl24_7 | ~ spl24_133)), inference(forward_demodulation, [], [f1300, f720])).
fof(f1300, plain, ((e4 = op(op(e4, e3), e3)) | ~ spl24_133), inference(avatar_component_clause, [], [f1298])).
fof(f17454, plain, (spl24_27 | ~ spl24_7 | ~ spl24_134), inference(avatar_split_clause, [], [f17453, f1303, f718, f802])).
fof(f17453, plain, ((e1 = op(e3, e4)) | (~ spl24_7 | ~ spl24_134)), inference(forward_demodulation, [], [f1304, f720])).
fof(f17432, plain, (~ spl24_45 | ~ spl24_225 | spl24_267), inference(avatar_split_clause, [], [f17427, f2005, f1746, f877])).
fof(f2005, plain, (spl24_267 <=> (e4 = op(e3, op(op(e3, op(e3, e3)), e3)))), introduced(avatar_definition, [new_symbols(naming, [spl24_267])])).
fof(f17427, plain, (~ (e4 = op(e3, e1)) | (~ spl24_225 | spl24_267)), inference(backward_demodulation, [], [f2007, f1747])).
fof(f1747, plain, ((e1 = op(op(e3, op(e3, e3)), e3)) | ~ spl24_225), inference(avatar_component_clause, [], [f1746])).
fof(f2007, plain, (~ (e4 = op(e3, op(op(e3, op(e3, e3)), e3))) | spl24_267), inference(avatar_component_clause, [], [f2005])).
fof(f17431, plain, (~ spl24_41 | spl24_210 | ~ spl24_225), inference(avatar_split_clause, [], [f17426, f1746, f1680, f861])).
fof(f1680, plain, (spl24_210 <=> (e0 = op(e3, op(op(e3, op(e3, e3)), e3)))), introduced(avatar_definition, [new_symbols(naming, [spl24_210])])).
fof(f17426, plain, (~ (e0 = op(e3, e1)) | (spl24_210 | ~ spl24_225)), inference(backward_demodulation, [], [f1682, f1747])).
fof(f1682, plain, (~ (e0 = op(e3, op(op(e3, op(e3, e3)), e3))) | spl24_210), inference(avatar_component_clause, [], [f1680])).
fof(f17414, plain, (~ spl24_92 | ~ spl24_97), inference(avatar_split_clause, [], [f17413, f1096, f1075])).
fof(f1096, plain, (spl24_97 <=> (e1 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_97])])).
fof(f17413, plain, (~ (e1 = op(e1, e1)) | ~ spl24_97), inference(forward_demodulation, [], [f442, f1098])).
fof(f1098, plain, ((e1 = op(e1, e0)) | ~ spl24_97), inference(avatar_component_clause, [], [f1096])).
fof(f442, plain, ~ (op(e1, e0) = op(e1, e1)), inference(cnf_transformation, [], [f4])).
fof(f17412, plain, (~ spl24_92 | ~ spl24_117), inference(avatar_split_clause, [], [f17411, f1180, f1075])).
fof(f1180, plain, (spl24_117 <=> (e1 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_117])])).
fof(f17411, plain, (~ (e1 = op(e1, e1)) | ~ spl24_117), inference(forward_demodulation, [], [f392, f1182])).
fof(f1182, plain, ((e1 = op(e0, e1)) | ~ spl24_117), inference(avatar_component_clause, [], [f1180])).
fof(f17400, plain, (~ spl24_20 | ~ spl24_25), inference(avatar_split_clause, [], [f17399, f793, f772])).
fof(f793, plain, (spl24_25 <=> (e4 = op(e4, e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_25])])).
fof(f17399, plain, (~ (e4 = op(e4, e1)) | ~ spl24_25), inference(forward_demodulation, [], [f472, f795])).
fof(f795, plain, ((e4 = op(e4, e0)) | ~ spl24_25), inference(avatar_component_clause, [], [f793])).
fof(f472, plain, ~ (op(e4, e0) = op(e4, e1)), inference(cnf_transformation, [], [f4])).
fof(f17388, plain, (~ spl24_14 | ~ spl24_37 | ~ spl24_137), inference(avatar_contradiction_clause, [], [f17387])).
fof(f17387, plain, ($false | (~ spl24_14 | ~ spl24_37 | ~ spl24_137)), inference(subsumption_resolution, [], [f17386, f488])).
fof(f488, plain, ~ (e1 = e4), inference(cnf_transformation, [], [f5])).
fof(f17386, plain, ((e1 = e4) | (~ spl24_14 | ~ spl24_37 | ~ spl24_137)), inference(forward_demodulation, [], [f17382, f846])).
fof(f17382, plain, ((e4 = op(e3, e2)) | (~ spl24_14 | ~ spl24_137)), inference(backward_demodulation, [], [f1319, f749])).
fof(f749, plain, ((e3 = op(e4, e2)) | ~ spl24_14), inference(avatar_component_clause, [], [f747])).
fof(f17385, plain, (~ spl24_14 | ~ spl24_28 | spl24_136), inference(avatar_contradiction_clause, [], [f17384])).
fof(f17384, plain, ($false | (~ spl24_14 | ~ spl24_28 | spl24_136)), inference(subsumption_resolution, [], [f17381, f808])).
fof(f17381, plain, (~ (e2 = op(e3, e4)) | (~ spl24_14 | spl24_136)), inference(backward_demodulation, [], [f1314, f749])).
fof(f1314, plain, (~ (e2 = op(op(e4, e2), e4)) | spl24_136), inference(avatar_component_clause, [], [f1312])).
fof(f1312, plain, (spl24_136 <=> (e2 = op(op(e4, e2), e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_136])])).
fof(f17383, plain, (~ spl24_4 | ~ spl24_14), inference(avatar_split_clause, [], [f17379, f747, f705])).
fof(f17379, plain, (~ (e3 = op(e4, e4)) | ~ spl24_14), inference(backward_demodulation, [], [f480, f749])).
fof(f17346, plain, (~ spl24_56 | ~ spl24_61), inference(avatar_split_clause, [], [f17333, f945, f924])).
fof(f17333, plain, (~ (e0 = op(e2, e3)) | ~ spl24_61), inference(backward_demodulation, [], [f459, f947])).
fof(f459, plain, ~ (op(e2, e2) = op(e2, e3)), inference(cnf_transformation, [], [f4])).
fof(f17345, plain, (~ spl24_11 | ~ spl24_61), inference(avatar_split_clause, [], [f17332, f945, f735])).
fof(f17332, plain, (~ (e0 = op(e4, e2)) | ~ spl24_61), inference(backward_demodulation, [], [f410, f947])).
fof(f410, plain, ~ (op(e2, e2) = op(e4, e2)), inference(cnf_transformation, [], [f4])).
fof(f17321, plain, (~ spl24_78 | ~ spl24_83), inference(avatar_split_clause, [], [f17311, f1037, f1016])).
fof(f1016, plain, (spl24_78 <=> (e2 = op(e1, e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_78])])).
fof(f17311, plain, (~ (e2 = op(e1, e4)) | ~ spl24_83), inference(backward_demodulation, [], [f451, f1039])).
fof(f1039, plain, ((e2 = op(e1, e3)) | ~ spl24_83), inference(avatar_component_clause, [], [f1037])).
fof(f17320, plain, (~ spl24_33 | ~ spl24_83), inference(avatar_split_clause, [], [f17310, f1037, f827])).
fof(f17310, plain, (~ (e2 = op(e3, e3)) | ~ spl24_83), inference(backward_demodulation, [], [f417, f1039])).
fof(f17308, plain, (~ spl24_85 | ~ spl24_90), inference(avatar_split_clause, [], [f17303, f1066, f1045])).
fof(f1066, plain, (spl24_90 <=> (e4 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_90])])).
fof(f17303, plain, (~ (e4 = op(e1, e3)) | ~ spl24_90), inference(backward_demodulation, [], [f449, f1068])).
fof(f1068, plain, ((e4 = op(e1, e2)) | ~ spl24_90), inference(avatar_component_clause, [], [f1066])).
fof(f449, plain, ~ (op(e1, e2) = op(e1, e3)), inference(cnf_transformation, [], [f4])).
fof(f17307, plain, (~ spl24_65 | ~ spl24_90), inference(avatar_split_clause, [], [f17301, f1066, f961])).
fof(f17301, plain, (~ (e4 = op(e2, e2)) | ~ spl24_90), inference(backward_demodulation, [], [f406, f1068])).
fof(f406, plain, ~ (op(e1, e2) = op(e2, e2)), inference(cnf_transformation, [], [f4])).
fof(f17296, plain, (~ spl24_78 | ~ spl24_18 | spl24_142), inference(avatar_split_clause, [], [f17295, f1341, f764, f1016])).
fof(f1341, plain, (spl24_142 <=> (op(e1, e4) = op(e4, e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_142])])).
fof(f17295, plain, (~ (e2 = op(e1, e4)) | (~ spl24_18 | spl24_142)), inference(forward_demodulation, [], [f1343, f766])).
fof(f1343, plain, (~ (op(e1, e4) = op(e4, e1)) | spl24_142), inference(avatar_component_clause, [], [f1341])).
fof(f17292, plain, (~ spl24_89 | ~ spl24_37 | spl24_154), inference(avatar_split_clause, [], [f17291, f1400, f844, f1062])).
fof(f17291, plain, (~ (e3 = op(e1, e2)) | (~ spl24_37 | spl24_154)), inference(forward_demodulation, [], [f1401, f846])).
fof(f1401, plain, (~ (e3 = op(op(e3, e2), e2)) | spl24_154), inference(avatar_component_clause, [], [f1400])).
fof(f17286, plain, (spl24_90 | ~ spl24_52 | ~ spl24_165), inference(avatar_split_clause, [], [f17285, f1452, f907, f1066])).
fof(f17285, plain, ((e4 = op(e1, e2)) | (~ spl24_52 | ~ spl24_165)), inference(forward_demodulation, [], [f1453, f909])).
fof(f909, plain, ((e1 = op(e2, e4)) | ~ spl24_52), inference(avatar_component_clause, [], [f907])).
fof(f1453, plain, ((e4 = op(op(e2, e4), e2)) | ~ spl24_165), inference(avatar_component_clause, [], [f1452])).
fof(f17276, plain, (~ spl24_85 | spl24_223 | ~ spl24_233), inference(avatar_split_clause, [], [f17275, f1783, f1737, f1045])).
fof(f1737, plain, (spl24_223 <=> (e4 = op(op(e3, op(e3, e3)), e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_223])])).
fof(f1783, plain, (spl24_233 <=> (e1 = op(e3, op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl24_233])])).
fof(f17275, plain, (~ (e4 = op(e1, e3)) | (spl24_223 | ~ spl24_233)), inference(forward_demodulation, [], [f1739, f1784])).
fof(f1784, plain, ((e1 = op(e3, op(e3, e3))) | ~ spl24_233), inference(avatar_component_clause, [], [f1783])).
fof(f1739, plain, (~ (e4 = op(op(e3, op(e3, e3)), e3)) | spl24_223), inference(avatar_component_clause, [], [f1737])).
fof(f17262, plain, (~ spl24_56 | spl24_216 | ~ spl24_217), inference(avatar_split_clause, [], [f17261, f1710, f1706, f924])).
fof(f1706, plain, (spl24_216 <=> (e0 = op(e2, op(op(e2, op(e2, e2)), e2)))), introduced(avatar_definition, [new_symbols(naming, [spl24_216])])).
fof(f1710, plain, (spl24_217 <=> (e3 = op(op(e2, op(e2, e2)), e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_217])])).
fof(f17261, plain, (~ (e0 = op(e2, e3)) | (spl24_216 | ~ spl24_217)), inference(backward_demodulation, [], [f1708, f1711])).
fof(f1711, plain, ((e3 = op(op(e2, op(e2, e2)), e2)) | ~ spl24_217), inference(avatar_component_clause, [], [f1710])).
fof(f1708, plain, (~ (e0 = op(e2, op(op(e2, op(e2, e2)), e2))) | spl24_216), inference(avatar_component_clause, [], [f1706])).
fof(f17244, plain, (~ spl24_83 | ~ spl24_37 | spl24_153), inference(avatar_split_clause, [], [f17208, f1395, f844, f1037])).
fof(f17208, plain, (~ (e2 = op(e1, e3)) | (~ spl24_37 | spl24_153)), inference(backward_demodulation, [], [f1397, f846])).
fof(f17168, plain, (~ spl24_11 | ~ spl24_78 | spl24_213 | ~ spl24_230), inference(avatar_split_clause, [], [f17164, f1769, f1693, f1016, f735])).
fof(f17164, plain, (~ (e0 = op(e4, e2)) | (~ spl24_78 | spl24_213 | ~ spl24_230)), inference(backward_demodulation, [], [f16937, f1018])).
fof(f1018, plain, ((e2 = op(e1, e4)) | ~ spl24_78), inference(avatar_component_clause, [], [f1016])).
fof(f16937, plain, (~ (e0 = op(e4, op(e1, e4))) | (spl24_213 | ~ spl24_230)), inference(backward_demodulation, [], [f1695, f1770])).
fof(f17167, plain, (spl24_52 | ~ spl24_78 | ~ spl24_182), inference(avatar_split_clause, [], [f17163, f1536, f1016, f907])).
fof(f1536, plain, (spl24_182 <=> (e1 = op(op(e1, e4), e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_182])])).
fof(f17163, plain, ((e1 = op(e2, e4)) | (~ spl24_78 | ~ spl24_182)), inference(backward_demodulation, [], [f1538, f1018])).
fof(f1538, plain, ((e1 = op(op(e1, e4), e4)) | ~ spl24_182), inference(avatar_component_clause, [], [f1536])).
fof(f17166, plain, (~ spl24_28 | ~ spl24_78), inference(avatar_split_clause, [], [f17158, f1016, f806])).
fof(f17158, plain, (~ (e2 = op(e3, e4)) | ~ spl24_78), inference(backward_demodulation, [], [f427, f1018])).
fof(f17154, plain, (spl24_37 | ~ spl24_89 | ~ spl24_188), inference(avatar_split_clause, [], [f17145, f1566, f1062, f844])).
fof(f1566, plain, (spl24_188 <=> (e1 = op(op(e1, e2), e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_188])])).
fof(f17145, plain, ((e1 = op(e3, e2)) | (~ spl24_89 | ~ spl24_188)), inference(backward_demodulation, [], [f1568, f1064])).
fof(f1568, plain, ((e1 = op(op(e1, e2), e2)) | ~ spl24_188), inference(avatar_component_clause, [], [f1566])).
fof(f17151, plain, (~ spl24_39 | ~ spl24_89), inference(avatar_split_clause, [], [f17139, f1062, f852])).
fof(f17139, plain, (~ (e3 = op(e3, e2)) | ~ spl24_89), inference(backward_demodulation, [], [f407, f1064])).
fof(f17130, plain, (~ spl24_55 | ~ spl24_105), inference(avatar_split_clause, [], [f17123, f1129, f919])).
fof(f17123, plain, (~ (e4 = op(e2, e4)) | ~ spl24_105), inference(backward_demodulation, [], [f423, f1131])).
fof(f1131, plain, ((e4 = op(e0, e4)) | ~ spl24_105), inference(avatar_component_clause, [], [f1129])).
fof(f17116, plain, (~ spl24_59 | ~ spl24_109), inference(avatar_split_clause, [], [f17106, f1146, f936])).
fof(f17106, plain, (~ (e3 = op(e2, e3)) | ~ spl24_109), inference(backward_demodulation, [], [f413, f1148])).
fof(f1148, plain, ((e3 = op(e0, e3)) | ~ spl24_109), inference(avatar_component_clause, [], [f1146])).
fof(f413, plain, ~ (op(e0, e3) = op(e2, e3)), inference(cnf_transformation, [], [f4])).
fof(f17103, plain, (~ spl24_108 | ~ spl24_113), inference(avatar_split_clause, [], [f17090, f1163, f1142])).
fof(f1163, plain, (spl24_113 <=> (e2 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_113])])).
fof(f17090, plain, (~ (e2 = op(e0, e3)) | ~ spl24_113), inference(backward_demodulation, [], [f439, f1165])).
fof(f1165, plain, ((e2 = op(e0, e2)) | ~ spl24_113), inference(avatar_component_clause, [], [f1163])).
fof(f17102, plain, (~ spl24_63 | ~ spl24_113), inference(avatar_split_clause, [], [f17087, f1163, f953])).
fof(f17087, plain, (~ (e2 = op(e2, e2)) | ~ spl24_113), inference(backward_demodulation, [], [f403, f1165])).
fof(f17080, plain, (~ spl24_67 | ~ spl24_117), inference(avatar_split_clause, [], [f17066, f1180, f970])).
fof(f17066, plain, (~ (e1 = op(e2, e1)) | ~ spl24_117), inference(backward_demodulation, [], [f393, f1182])).
fof(f393, plain, ~ (op(e0, e1) = op(e2, e1)), inference(cnf_transformation, [], [f4])).
fof(f17043, plain, (~ spl24_111 | ~ spl24_121), inference(avatar_split_clause, [], [f17024, f1197, f1155])).
fof(f1197, plain, (spl24_121 <=> (e0 = op(e0, e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_121])])).
fof(f17024, plain, (~ (e0 = op(e0, e2)) | ~ spl24_121), inference(backward_demodulation, [], [f433, f1199])).
fof(f1199, plain, ((e0 = op(e0, e0)) | ~ spl24_121), inference(avatar_component_clause, [], [f1197])).
fof(f17042, plain, (~ spl24_71 | ~ spl24_121), inference(avatar_split_clause, [], [f17020, f1197, f987])).
fof(f17020, plain, (~ (e0 = op(e2, e0)) | ~ spl24_121), inference(backward_demodulation, [], [f383, f1199])).
fof(f17018, plain, (spl24_25 | ~ spl24_126), inference(avatar_split_clause, [], [f17004, f1218, f793])).
fof(f1218, plain, (spl24_126 <=> (e0 = unit)), introduced(avatar_definition, [new_symbols(naming, [spl24_126])])).
fof(f17004, plain, ((e4 = op(e4, e0)) | ~ spl24_126), inference(backward_demodulation, [], [f330, f1220])).
fof(f1220, plain, ((e0 = unit) | ~ spl24_126), inference(avatar_component_clause, [], [f1218])).
fof(f17017, plain, (spl24_105 | ~ spl24_126), inference(avatar_split_clause, [], [f17003, f1218, f1129])).
fof(f17003, plain, ((e4 = op(e0, e4)) | ~ spl24_126), inference(backward_demodulation, [], [f329, f1220])).
fof(f17016, plain, (spl24_49 | ~ spl24_126), inference(avatar_split_clause, [], [f17002, f1218, f894])).
fof(f894, plain, (spl24_49 <=> (e3 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_49])])).
fof(f17002, plain, ((e3 = op(e3, e0)) | ~ spl24_126), inference(backward_demodulation, [], [f328, f1220])).
fof(f17015, plain, (spl24_109 | ~ spl24_126), inference(avatar_split_clause, [], [f17001, f1218, f1146])).
fof(f17001, plain, ((e3 = op(e0, e3)) | ~ spl24_126), inference(backward_demodulation, [], [f327, f1220])).
fof(f17014, plain, (spl24_73 | ~ spl24_126), inference(avatar_split_clause, [], [f17000, f1218, f995])).
fof(f17000, plain, ((e2 = op(e2, e0)) | ~ spl24_126), inference(backward_demodulation, [], [f326, f1220])).
fof(f17013, plain, (spl24_113 | ~ spl24_126), inference(avatar_split_clause, [], [f16999, f1218, f1163])).
fof(f16999, plain, ((e2 = op(e0, e2)) | ~ spl24_126), inference(backward_demodulation, [], [f325, f1220])).
fof(f17012, plain, (spl24_97 | ~ spl24_126), inference(avatar_split_clause, [], [f16998, f1218, f1096])).
fof(f16998, plain, ((e1 = op(e1, e0)) | ~ spl24_126), inference(backward_demodulation, [], [f324, f1220])).
fof(f17011, plain, (spl24_117 | ~ spl24_126), inference(avatar_split_clause, [], [f16997, f1218, f1180])).
fof(f16997, plain, ((e1 = op(e0, e1)) | ~ spl24_126), inference(backward_demodulation, [], [f323, f1220])).
fof(f17010, plain, (spl24_121 | ~ spl24_126), inference(avatar_split_clause, [], [f16996, f1218, f1197])).
fof(f16996, plain, ((e0 = op(e0, e0)) | ~ spl24_126), inference(backward_demodulation, [], [f322, f1220])).
fof(f16953, plain, (~ spl24_26 | ~ spl24_85 | spl24_210 | ~ spl24_233), inference(avatar_split_clause, [], [f16952, f1783, f1680, f1045, f798])).
fof(f16952, plain, (~ (e0 = op(e3, e4)) | (~ spl24_85 | spl24_210 | ~ spl24_233)), inference(forward_demodulation, [], [f16951, f1047])).
fof(f16951, plain, (~ (e0 = op(e3, op(e1, e3))) | (spl24_210 | ~ spl24_233)), inference(forward_demodulation, [], [f1682, f1784])).
fof(f16941, plain, (~ spl24_79 | spl24_219 | ~ spl24_230), inference(avatar_split_clause, [], [f16938, f1769, f1719, f1020])).
fof(f1719, plain, (spl24_219 <=> (e3 = op(op(e4, op(e4, e4)), e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_219])])).
fof(f16938, plain, (~ (e3 = op(e1, e4)) | (spl24_219 | ~ spl24_230)), inference(backward_demodulation, [], [f1721, f1770])).
fof(f1721, plain, (~ (e3 = op(op(e4, op(e4, e4)), e4)) | spl24_219), inference(avatar_component_clause, [], [f1719])).
fof(f16935, plain, (~ spl24_98 | ~ spl24_91 | spl24_236), inference(avatar_split_clause, [], [f16934, f1800, f1071, f1100])).
fof(f1071, plain, (spl24_91 <=> (e0 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_91])])).
fof(f1800, plain, (spl24_236 <=> (e2 = op(e1, op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl24_236])])).
fof(f16934, plain, (~ (e2 = op(e1, e0)) | (~ spl24_91 | spl24_236)), inference(forward_demodulation, [], [f1802, f1073])).
fof(f1073, plain, ((e0 = op(e1, e1)) | ~ spl24_91), inference(avatar_component_clause, [], [f1071])).
fof(f1802, plain, (~ (e2 = op(e1, op(e1, e1))) | spl24_236), inference(avatar_component_clause, [], [f1800])).
fof(f16890, plain, (~ spl24_35 | ~ spl24_85), inference(avatar_split_clause, [], [f16889, f1045, f835])).
fof(f835, plain, (spl24_35 <=> (e4 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_35])])).
fof(f16889, plain, (~ (e4 = op(e3, e3)) | ~ spl24_85), inference(forward_demodulation, [], [f417, f1047])).
fof(f16887, plain, (~ spl24_28 | ~ spl24_223 | spl24_254), inference(avatar_split_clause, [], [f16886, f1904, f1737, f806])).
fof(f1904, plain, (spl24_254 <=> (e2 = op(e3, op(op(e3, op(e3, e3)), e3)))), introduced(avatar_definition, [new_symbols(naming, [spl24_254])])).
fof(f16886, plain, (~ (e2 = op(e3, e4)) | (~ spl24_223 | spl24_254)), inference(forward_demodulation, [], [f1906, f1738])).
fof(f1738, plain, ((e4 = op(op(e3, op(e3, e3)), e3)) | ~ spl24_223), inference(avatar_component_clause, [], [f1737])).
fof(f1906, plain, (~ (e2 = op(e3, op(op(e3, op(e3, e3)), e3))) | spl24_254), inference(avatar_component_clause, [], [f1904])).
fof(f16883, plain, (~ spl24_30 | ~ spl24_45), inference(avatar_split_clause, [], [f16882, f877, f814])).
fof(f16882, plain, (~ (e4 = op(e3, e4)) | ~ spl24_45), inference(forward_demodulation, [], [f468, f879])).
fof(f879, plain, ((e4 = op(e3, e1)) | ~ spl24_45), inference(avatar_component_clause, [], [f877])).
fof(f468, plain, ~ (op(e3, e1) = op(e3, e4)), inference(cnf_transformation, [], [f4])).
fof(f16867, plain, (~ spl24_24 | ~ spl24_47 | ~ spl24_145), inference(avatar_contradiction_clause, [], [f16866])).
fof(f16866, plain, ($false | (~ spl24_24 | ~ spl24_47 | ~ spl24_145)), inference(subsumption_resolution, [], [f16865, f488])).
fof(f16865, plain, ((e1 = e4) | (~ spl24_24 | ~ spl24_47 | ~ spl24_145)), inference(forward_demodulation, [], [f16863, f888])).
fof(f16863, plain, ((e4 = op(e3, e0)) | (~ spl24_24 | ~ spl24_145)), inference(backward_demodulation, [], [f1357, f791])).
fof(f16852, plain, (~ spl24_28 | ~ spl24_31 | ~ spl24_47 | ~ spl24_85 | spl24_254), inference(avatar_contradiction_clause, [], [f16851])).
fof(f16851, plain, ($false | (~ spl24_28 | ~ spl24_31 | ~ spl24_47 | ~ spl24_85 | spl24_254)), inference(subsumption_resolution, [], [f16848, f808])).
fof(f16848, plain, (~ (e2 = op(e3, e4)) | (~ spl24_31 | ~ spl24_47 | ~ spl24_85 | spl24_254)), inference(backward_demodulation, [], [f16767, f1047])).
fof(f16767, plain, (~ (e2 = op(e3, op(e1, e3))) | (~ spl24_31 | ~ spl24_47 | spl24_254)), inference(forward_demodulation, [], [f16766, f888])).
fof(f16766, plain, (~ (e2 = op(e3, op(op(e3, e0), e3))) | (~ spl24_31 | spl24_254)), inference(forward_demodulation, [], [f1906, f821])).
fof(f16843, plain, (~ spl24_76 | ~ spl24_91), inference(avatar_split_clause, [], [f16835, f1071, f1008])).
fof(f16835, plain, (~ (e0 = op(e1, e4)) | ~ spl24_91), inference(backward_demodulation, [], [f448, f1073])).
fof(f448, plain, ~ (op(e1, e1) = op(e1, e4)), inference(cnf_transformation, [], [f4])).
fof(f16842, plain, (~ spl24_16 | ~ spl24_91), inference(avatar_split_clause, [], [f16833, f1071, f756])).
fof(f16833, plain, (~ (e0 = op(e4, e1)) | ~ spl24_91), inference(backward_demodulation, [], [f398, f1073])).
fof(f16827, plain, (~ spl24_83 | ~ spl24_98), inference(avatar_split_clause, [], [f16816, f1100, f1037])).
fof(f16816, plain, (~ (e2 = op(e1, e3)) | ~ spl24_98), inference(backward_demodulation, [], [f444, f1102])).
fof(f16826, plain, (~ spl24_23 | ~ spl24_98), inference(avatar_split_clause, [], [f16814, f1100, f785])).
fof(f16814, plain, (~ (e2 = op(e4, e0)) | ~ spl24_98), inference(backward_demodulation, [], [f388, f1102])).
fof(f388, plain, ~ (op(e1, e0) = op(e4, e0)), inference(cnf_transformation, [], [f4])).
fof(f16812, plain, (~ spl24_83 | ~ spl24_108), inference(avatar_split_clause, [], [f16807, f1142, f1037])).
fof(f16807, plain, (~ (e2 = op(e1, e3)) | ~ spl24_108), inference(backward_demodulation, [], [f412, f1144])).
fof(f16798, plain, (~ spl24_110 | ~ spl24_125), inference(avatar_split_clause, [], [f16795, f1213, f1150])).
fof(f1213, plain, (spl24_125 <=> (op(e0, e0) = e4)), introduced(avatar_definition, [new_symbols(naming, [spl24_125])])).
fof(f16795, plain, (~ (e4 = op(e0, e3)) | ~ spl24_125), inference(backward_demodulation, [], [f434, f1215])).
fof(f1215, plain, ((op(e0, e0) = e4) | ~ spl24_125), inference(avatar_component_clause, [], [f1213])).
fof(f16797, plain, (~ spl24_100 | ~ spl24_125), inference(avatar_split_clause, [], [f16792, f1213, f1108])).
fof(f16792, plain, (~ (e4 = op(e1, e0)) | ~ spl24_125), inference(backward_demodulation, [], [f382, f1215])).
fof(f16787, plain, (~ spl24_100 | ~ spl24_102 | spl24_196), inference(avatar_split_clause, [], [f16786, f1606, f1117, f1108])).
fof(f1606, plain, (spl24_196 <=> (e4 = op(op(e0, e4), e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_196])])).
fof(f16786, plain, (~ (e4 = op(e1, e0)) | (~ spl24_102 | spl24_196)), inference(forward_demodulation, [], [f1608, f1119])).
fof(f1608, plain, (~ (e4 = op(op(e0, e4), e0)) | spl24_196), inference(avatar_component_clause, [], [f1606])).
fof(f16727, plain, (~ spl24_77 | ~ spl24_102), inference(avatar_split_clause, [], [f16726, f1117, f1012])).
fof(f1012, plain, (spl24_77 <=> (e1 = op(e1, e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_77])])).
fof(f16726, plain, (~ (e1 = op(e1, e4)) | ~ spl24_102), inference(forward_demodulation, [], [f422, f1119])).
fof(f422, plain, ~ (op(e0, e4) = op(e1, e4)), inference(cnf_transformation, [], [f4])).
fof(f16721, plain, (~ spl24_4 | ~ spl24_7 | spl24_230), inference(avatar_contradiction_clause, [], [f16720])).
fof(f16720, plain, ($false | (~ spl24_4 | ~ spl24_7 | spl24_230)), inference(subsumption_resolution, [], [f16715, f720])).
fof(f16715, plain, (~ (e1 = op(e4, e3)) | (~ spl24_4 | spl24_230)), inference(backward_demodulation, [], [f1771, f707])).
fof(f707, plain, ((e3 = op(e4, e4)) | ~ spl24_4), inference(avatar_component_clause, [], [f705])).
fof(f1771, plain, (~ (e1 = op(e4, op(e4, e4))) | spl24_230), inference(avatar_component_clause, [], [f1769])).
fof(f16712, plain, (~ spl24_7 | ~ spl24_83 | ~ spl24_133), inference(avatar_contradiction_clause, [], [f16711])).
fof(f16711, plain, ($false | (~ spl24_7 | ~ spl24_83 | ~ spl24_133)), inference(subsumption_resolution, [], [f16710, f490])).
fof(f490, plain, ~ (e2 = e4), inference(cnf_transformation, [], [f5])).
fof(f16710, plain, ((e2 = e4) | (~ spl24_7 | ~ spl24_83 | ~ spl24_133)), inference(forward_demodulation, [], [f16707, f1039])).
fof(f16707, plain, ((e4 = op(e1, e3)) | (~ spl24_7 | ~ spl24_133)), inference(backward_demodulation, [], [f1300, f720])).
fof(f16708, plain, (~ spl24_2 | ~ spl24_7), inference(avatar_split_clause, [], [f16702, f718, f697])).
fof(f697, plain, (spl24_2 <=> (e1 = op(e4, e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_2])])).
fof(f16702, plain, (~ (e1 = op(e4, e4)) | ~ spl24_7), inference(backward_demodulation, [], [f481, f720])).
fof(f16700, plain, (~ spl24_6 | ~ spl24_16), inference(avatar_split_clause, [], [f16692, f756, f714])).
fof(f16692, plain, (~ (e0 = op(e4, e3)) | ~ spl24_16), inference(backward_demodulation, [], [f477, f758])).
fof(f758, plain, ((e0 = op(e4, e1)) | ~ spl24_16), inference(avatar_component_clause, [], [f756])).
fof(f16691, plain, (~ spl24_6 | ~ spl24_31), inference(avatar_split_clause, [], [f16684, f819, f714])).
fof(f16684, plain, (~ (e0 = op(e4, e3)) | ~ spl24_31), inference(backward_demodulation, [], [f421, f821])).
fof(f16683, plain, (~ spl24_19 | ~ spl24_45 | spl24_158), inference(avatar_split_clause, [], [f16680, f1419, f877, f768])).
fof(f1419, plain, (spl24_158 <=> (e3 = op(op(e3, e1), e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_158])])).
fof(f16680, plain, (~ (e3 = op(e4, e1)) | (~ spl24_45 | spl24_158)), inference(backward_demodulation, [], [f1420, f879])).
fof(f1420, plain, (~ (e3 = op(op(e3, e1), e1)) | spl24_158), inference(avatar_component_clause, [], [f1419])).
fof(f16681, plain, (~ spl24_35 | ~ spl24_45), inference(avatar_split_clause, [], [f16678, f877, f835])).
fof(f16678, plain, (~ (e4 = op(e3, e3)) | ~ spl24_45), inference(backward_demodulation, [], [f467, f879])).
fof(f16676, plain, (~ spl24_19 | ~ spl24_94), inference(avatar_split_clause, [], [f16672, f1083, f768])).
fof(f16672, plain, (~ (e3 = op(e4, e1)) | ~ spl24_94), inference(backward_demodulation, [], [f398, f1085])).
fof(f16669, plain, (~ spl24_95 | ~ spl24_100), inference(avatar_split_clause, [], [f16666, f1108, f1087])).
fof(f1087, plain, (spl24_95 <=> (e4 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_95])])).
fof(f16666, plain, (~ (e4 = op(e1, e1)) | ~ spl24_100), inference(backward_demodulation, [], [f442, f1110])).
fof(f16664, plain, (~ spl24_2 | ~ spl24_102), inference(avatar_split_clause, [], [f16658, f1117, f697])).
fof(f16658, plain, (~ (e1 = op(e4, e4)) | ~ spl24_102), inference(backward_demodulation, [], [f425, f1119])).
fof(f16657, plain, (~ spl24_6 | ~ spl24_110 | spl24_200), inference(avatar_split_clause, [], [f16654, f1626, f1150, f714])).
fof(f16654, plain, (~ (e0 = op(e4, e3)) | (~ spl24_110 | spl24_200)), inference(backward_demodulation, [], [f1627, f1152])).
fof(f1627, plain, (~ (e0 = op(op(e0, e3), e3)) | spl24_200), inference(avatar_component_clause, [], [f1626])).
fof(f16655, plain, (~ spl24_35 | ~ spl24_110), inference(avatar_split_clause, [], [f16650, f1150, f835])).
fof(f16650, plain, (~ (e4 = op(e3, e3)) | ~ spl24_110), inference(backward_demodulation, [], [f414, f1152])).
fof(f414, plain, ~ (op(e0, e3) = op(e3, e3)), inference(cnf_transformation, [], [f4])).
fof(f16648, plain, (spl24_110 | ~ spl24_124 | ~ spl24_245), inference(avatar_split_clause, [], [f16642, f1852, f1209, f1150])).
fof(f16642, plain, ((e4 = op(e0, e3)) | (~ spl24_124 | ~ spl24_245)), inference(backward_demodulation, [], [f1853, f1211])).
fof(f16646, plain, (~ spl24_99 | ~ spl24_124), inference(avatar_split_clause, [], [f16638, f1209, f1104])).
fof(f1104, plain, (spl24_99 <=> (e3 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_99])])).
fof(f16638, plain, (~ (e3 = op(e1, e0)) | ~ spl24_124), inference(backward_demodulation, [], [f382, f1211])).
fof(f16637, plain, (spl24_16 | ~ spl24_76 | ~ spl24_142), inference(avatar_split_clause, [], [f16636, f1341, f1008, f756])).
fof(f16636, plain, ((e0 = op(e4, e1)) | (~ spl24_76 | ~ spl24_142)), inference(forward_demodulation, [], [f1342, f1010])).
fof(f1342, plain, ((op(e1, e4) = op(e4, e1)) | ~ spl24_142), inference(avatar_component_clause, [], [f1341])).
fof(f16633, plain, (~ spl24_107 | ~ spl24_47 | spl24_163), inference(avatar_split_clause, [], [f16632, f1443, f886, f1138])).
fof(f16632, plain, (~ (e1 = op(e0, e3)) | (~ spl24_47 | spl24_163)), inference(forward_demodulation, [], [f1445, f888])).
fof(f16630, plain, (spl24_102 | ~ spl24_76 | ~ spl24_182), inference(avatar_split_clause, [], [f16629, f1536, f1008, f1117])).
fof(f16629, plain, ((e1 = op(e0, e4)) | (~ spl24_76 | ~ spl24_182)), inference(forward_demodulation, [], [f1538, f1010])).
fof(f16543, plain, (~ spl24_19 | ~ spl24_41 | ~ spl24_141), inference(avatar_contradiction_clause, [], [f16542])).
fof(f16542, plain, ($false | (~ spl24_19 | ~ spl24_41 | ~ spl24_141)), inference(subsumption_resolution, [], [f16541, f485])).
fof(f16541, plain, ((e0 = e4) | (~ spl24_19 | ~ spl24_41 | ~ spl24_141)), inference(forward_demodulation, [], [f16540, f863])).
fof(f16540, plain, ((e4 = op(e3, e1)) | (~ spl24_19 | ~ spl24_141)), inference(backward_demodulation, [], [f1338, f770])).
fof(f1338, plain, ((e4 = op(op(e4, e1), e1)) | ~ spl24_141), inference(avatar_component_clause, [], [f1336])).
fof(f1336, plain, (spl24_141 <=> (e4 = op(op(e4, e1), e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_141])])).
fof(f16534, plain, (~ spl24_3 | ~ spl24_23), inference(avatar_split_clause, [], [f16529, f785, f701])).
fof(f16529, plain, (~ (e2 = op(e4, e4)) | ~ spl24_23), inference(backward_demodulation, [], [f475, f787])).
fof(f16514, plain, (spl24_28 | ~ spl24_35 | ~ spl24_224), inference(avatar_split_clause, [], [f16509, f1741, f835, f806])).
fof(f1741, plain, (spl24_224 <=> (e2 = op(e3, op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl24_224])])).
fof(f16509, plain, ((e2 = op(e3, e4)) | (~ spl24_35 | ~ spl24_224)), inference(backward_demodulation, [], [f1742, f837])).
fof(f837, plain, ((e4 = op(e3, e3)) | ~ spl24_35), inference(avatar_component_clause, [], [f835])).
fof(f1742, plain, ((e2 = op(e3, op(e3, e3))) | ~ spl24_224), inference(avatar_component_clause, [], [f1741])).
fof(f16504, plain, (~ spl24_16 | ~ spl24_41), inference(avatar_split_clause, [], [f16495, f861, f756])).
fof(f16495, plain, (~ (e0 = op(e4, e1)) | ~ spl24_41), inference(backward_demodulation, [], [f401, f863])).
fof(f401, plain, ~ (op(e3, e1) = op(e4, e1)), inference(cnf_transformation, [], [f4])).
fof(f16485, plain, (~ spl24_30 | ~ spl24_55), inference(avatar_split_clause, [], [f16477, f919, f814])).
fof(f16477, plain, (~ (e4 = op(e3, e4)) | ~ spl24_55), inference(backward_demodulation, [], [f429, f921])).
fof(f921, plain, ((e4 = op(e2, e4)) | ~ spl24_55), inference(avatar_component_clause, [], [f919])).
fof(f429, plain, ~ (op(e2, e4) = op(e3, e4)), inference(cnf_transformation, [], [f4])).
fof(f16474, plain, (~ spl24_34 | ~ spl24_59), inference(avatar_split_clause, [], [f16464, f936, f831])).
fof(f16464, plain, (~ (e3 = op(e3, e3)) | ~ spl24_59), inference(backward_demodulation, [], [f419, f938])).
fof(f938, plain, ((e3 = op(e2, e3)) | ~ spl24_59), inference(avatar_component_clause, [], [f936])).
fof(f16454, plain, (~ spl24_58 | ~ spl24_63), inference(avatar_split_clause, [], [f16442, f953, f932])).
fof(f16442, plain, (~ (e2 = op(e2, e3)) | ~ spl24_63), inference(backward_demodulation, [], [f459, f955])).
fof(f955, plain, ((e2 = op(e2, e2)) | ~ spl24_63), inference(avatar_component_clause, [], [f953])).
fof(f16432, plain, (~ spl24_16 | ~ spl24_76 | spl24_142), inference(avatar_split_clause, [], [f16427, f1341, f1008, f756])).
fof(f16427, plain, (~ (e0 = op(e4, e1)) | (~ spl24_76 | spl24_142)), inference(backward_demodulation, [], [f1343, f1010])).
fof(f16419, plain, (~ spl24_69 | ~ spl24_83 | spl24_184), inference(avatar_split_clause, [], [f16416, f1546, f1037, f978])).
fof(f1546, plain, (spl24_184 <=> (e3 = op(op(e1, e3), e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_184])])).
fof(f16416, plain, (~ (e3 = op(e2, e1)) | (~ spl24_83 | spl24_184)), inference(backward_demodulation, [], [f1548, f1039])).
fof(f1548, plain, (~ (e3 = op(op(e1, e3), e1)) | spl24_184), inference(avatar_component_clause, [], [f1546])).
fof(f16418, plain, (~ spl24_58 | ~ spl24_83), inference(avatar_split_clause, [], [f16409, f1037, f932])).
fof(f16409, plain, (~ (e2 = op(e2, e3)) | ~ spl24_83), inference(backward_demodulation, [], [f416, f1039])).
fof(f16406, plain, (~ spl24_82 | ~ spl24_87), inference(avatar_split_clause, [], [f16393, f1054, f1033])).
fof(f16393, plain, (~ (e1 = op(e1, e3)) | ~ spl24_87), inference(backward_demodulation, [], [f449, f1056])).
fof(f16405, plain, (~ spl24_12 | ~ spl24_87), inference(avatar_split_clause, [], [f16392, f1054, f739])).
fof(f16392, plain, (~ (e1 = op(e4, e2)) | ~ spl24_87), inference(backward_demodulation, [], [f408, f1056])).
fof(f16389, plain, (spl24_47 | ~ spl24_99 | ~ spl24_193), inference(avatar_split_clause, [], [f16385, f1592, f1104, f886])).
fof(f16385, plain, ((e1 = op(e3, e0)) | (~ spl24_99 | ~ spl24_193)), inference(backward_demodulation, [], [f1594, f1106])).
fof(f1106, plain, ((e3 = op(e1, e0)) | ~ spl24_99), inference(avatar_component_clause, [], [f1104])).
fof(f1594, plain, ((e1 = op(op(e1, e0), e0)) | ~ spl24_193), inference(avatar_component_clause, [], [f1592])).
fof(f16388, plain, (spl24_41 | ~ spl24_99 | ~ spl24_192), inference(avatar_split_clause, [], [f16384, f1587, f1104, f861])).
fof(f16384, plain, ((e0 = op(e3, e1)) | (~ spl24_99 | ~ spl24_192)), inference(backward_demodulation, [], [f1588, f1106])).
fof(f16387, plain, (~ spl24_79 | ~ spl24_99), inference(avatar_split_clause, [], [f16382, f1104, f1020])).
fof(f16382, plain, (~ (e3 = op(e1, e4)) | ~ spl24_99), inference(backward_demodulation, [], [f445, f1106])).
fof(f16386, plain, (~ spl24_24 | ~ spl24_99), inference(avatar_split_clause, [], [f16379, f1104, f789])).
fof(f16379, plain, (~ (e3 = op(e4, e0)) | ~ spl24_99), inference(backward_demodulation, [], [f388, f1106])).
fof(f16378, plain, (~ spl24_47 | ~ spl24_104 | ~ spl24_125 | spl24_259), inference(avatar_split_clause, [], [f16373, f1938, f1213, f1125, f886])).
fof(f16373, plain, (~ (e1 = op(e3, e0)) | (~ spl24_104 | ~ spl24_125 | spl24_259)), inference(backward_demodulation, [], [f16238, f1127])).
fof(f16238, plain, (~ (e1 = op(op(e0, e4), e0)) | (~ spl24_125 | spl24_259)), inference(forward_demodulation, [], [f1940, f1215])).
fof(f16376, plain, (~ spl24_79 | ~ spl24_104), inference(avatar_split_clause, [], [f16364, f1125, f1020])).
fof(f16364, plain, (~ (e3 = op(e1, e4)) | ~ spl24_104), inference(backward_demodulation, [], [f422, f1127])).
fof(f16361, plain, (spl24_47 | ~ spl24_107 | ~ spl24_163), inference(avatar_split_clause, [], [f16351, f1443, f1138, f886])).
fof(f16351, plain, ((e1 = op(e3, e0)) | (~ spl24_107 | ~ spl24_163)), inference(backward_demodulation, [], [f1444, f1140])).
fof(f1444, plain, ((op(e0, e3) = op(e3, e0)) | ~ spl24_163), inference(avatar_component_clause, [], [f1443])).
fof(f16359, plain, (~ spl24_82 | ~ spl24_107), inference(avatar_split_clause, [], [f16343, f1138, f1033])).
fof(f16343, plain, (~ (e1 = op(e1, e3)) | ~ spl24_107), inference(backward_demodulation, [], [f412, f1140])).
fof(f16338, plain, (~ spl24_106 | ~ spl24_111), inference(avatar_split_clause, [], [f16323, f1155, f1134])).
fof(f16323, plain, (~ (e0 = op(e0, e3)) | ~ spl24_111), inference(backward_demodulation, [], [f439, f1157])).
fof(f16337, plain, (~ spl24_86 | ~ spl24_111), inference(avatar_split_clause, [], [f16319, f1155, f1050])).
fof(f16319, plain, (~ (e0 = op(e1, e2)) | ~ spl24_111), inference(backward_demodulation, [], [f402, f1157])).
fof(f16270, plain, (~ spl24_98 | ~ spl24_118 | spl24_194), inference(avatar_split_clause, [], [f16269, f1597, f1184, f1100])).
fof(f16269, plain, (~ (e2 = op(e1, e0)) | (~ spl24_118 | spl24_194)), inference(forward_demodulation, [], [f1599, f1186])).
fof(f1186, plain, ((e2 = op(e0, e1)) | ~ spl24_118), inference(avatar_component_clause, [], [f1184])).
fof(f1599, plain, (~ (op(e0, e1) = op(e1, e0)) | spl24_194), inference(avatar_component_clause, [], [f1597])).
fof(f16268, plain, (~ spl24_72 | ~ spl24_118 | spl24_205), inference(avatar_split_clause, [], [f16267, f1651, f1184, f991])).
fof(f16267, plain, (~ (e1 = op(e2, e0)) | (~ spl24_118 | spl24_205)), inference(forward_demodulation, [], [f1653, f1186])).
fof(f1653, plain, (~ (e1 = op(op(e0, e1), e0)) | spl24_205), inference(avatar_component_clause, [], [f1651])).
fof(f16256, plain, (~ spl24_79 | ~ spl24_95 | spl24_232), inference(avatar_split_clause, [], [f16255, f1778, f1087, f1020])).
fof(f16255, plain, (~ (e3 = op(e1, e4)) | (~ spl24_95 | spl24_232)), inference(forward_demodulation, [], [f1780, f1089])).
fof(f1089, plain, ((e4 = op(e1, e1)) | ~ spl24_95), inference(avatar_component_clause, [], [f1087])).
fof(f1780, plain, (~ (e3 = op(e1, op(e1, e1))) | spl24_232), inference(avatar_component_clause, [], [f1778])).
fof(f16248, plain, (spl24_104 | ~ spl24_125 | ~ spl24_248), inference(avatar_split_clause, [], [f16247, f1866, f1213, f1125])).
fof(f16247, plain, ((e3 = op(e0, e4)) | (~ spl24_125 | ~ spl24_248)), inference(forward_demodulation, [], [f1867, f1215])).
fof(f16242, plain, (~ spl24_102 | ~ spl24_125 | spl24_260), inference(avatar_split_clause, [], [f16241, f1944, f1213, f1117])).
fof(f16241, plain, (~ (e1 = op(e0, e4)) | (~ spl24_125 | spl24_260)), inference(forward_demodulation, [], [f1946, f1215])).
fof(f16240, plain, (spl24_76 | ~ spl24_95 | ~ spl24_261), inference(avatar_split_clause, [], [f16239, f1949, f1087, f1008])).
fof(f1949, plain, (spl24_261 <=> (e0 = op(e1, op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl24_261])])).
fof(f16239, plain, ((e0 = op(e1, e4)) | (~ spl24_95 | ~ spl24_261)), inference(forward_demodulation, [], [f1950, f1089])).
fof(f1950, plain, ((e0 = op(e1, op(e1, e1))) | ~ spl24_261), inference(avatar_component_clause, [], [f1949])).
fof(f16227, plain, (~ spl24_50 | ~ spl24_125), inference(avatar_split_clause, [], [f16226, f1213, f898])).
fof(f16226, plain, (~ (e4 = op(e3, e0)) | ~ spl24_125), inference(forward_demodulation, [], [f384, f1215])).
fof(f384, plain, ~ (op(e0, e0) = op(e3, e0)), inference(cnf_transformation, [], [f4])).
fof(f16220, plain, (~ spl24_1 | ~ spl24_16), inference(avatar_split_clause, [], [f16213, f756, f693])).
fof(f16213, plain, (~ (e0 = op(e4, e4)) | ~ spl24_16), inference(backward_demodulation, [], [f478, f758])).
fof(f16211, plain, (~ spl24_19 | ~ spl24_24), inference(avatar_split_clause, [], [f16207, f789, f768])).
fof(f16207, plain, (~ (e3 = op(e4, e1)) | ~ spl24_24), inference(backward_demodulation, [], [f472, f791])).
fof(f16205, plain, (~ spl24_1 | ~ spl24_51), inference(avatar_split_clause, [], [f16196, f903, f693])).
fof(f16196, plain, (~ (e0 = op(e4, e4)) | ~ spl24_51), inference(backward_demodulation, [], [f430, f905])).
fof(f905, plain, ((e0 = op(e2, e4)) | ~ spl24_51), inference(avatar_component_clause, [], [f903])).
fof(f430, plain, ~ (op(e2, e4) = op(e4, e4)), inference(cnf_transformation, [], [f4])).
fof(f16194, plain, (~ spl24_54 | ~ spl24_69), inference(avatar_split_clause, [], [f16189, f978, f915])).
fof(f16189, plain, (~ (e3 = op(e2, e4)) | ~ spl24_69), inference(backward_demodulation, [], [f458, f980])).
fof(f980, plain, ((e3 = op(e2, e1)) | ~ spl24_69), inference(avatar_component_clause, [], [f978])).
fof(f16193, plain, (~ spl24_19 | ~ spl24_69), inference(avatar_split_clause, [], [f16188, f978, f768])).
fof(f16188, plain, (~ (e3 = op(e4, e1)) | ~ spl24_69), inference(backward_demodulation, [], [f400, f980])).
fof(f400, plain, ~ (op(e2, e1) = op(e4, e1)), inference(cnf_transformation, [], [f4])).
fof(f16180, plain, (~ spl24_19 | ~ spl24_79 | spl24_142), inference(avatar_split_clause, [], [f16171, f1341, f1020, f768])).
fof(f16171, plain, (~ (e3 = op(e4, e1)) | (~ spl24_79 | spl24_142)), inference(backward_demodulation, [], [f1343, f1022])).
fof(f16179, plain, (~ spl24_54 | ~ spl24_79), inference(avatar_split_clause, [], [f16168, f1020, f915])).
fof(f16168, plain, (~ (e3 = op(e2, e4)) | ~ spl24_79), inference(backward_demodulation, [], [f426, f1022])).
fof(f16166, plain, (~ spl24_98 | ~ spl24_102 | ~ spl24_114 | ~ spl24_125 | spl24_265), inference(avatar_contradiction_clause, [], [f16165])).
fof(f16165, plain, ($false | (~ spl24_98 | ~ spl24_102 | ~ spl24_114 | ~ spl24_125 | spl24_265)), inference(subsumption_resolution, [], [f16159, f1169])).
fof(f16159, plain, (~ (e3 = op(e0, e2)) | (~ spl24_98 | ~ spl24_102 | ~ spl24_125 | spl24_265)), inference(backward_demodulation, [], [f16100, f1102])).
fof(f16100, plain, (~ (e3 = op(e0, op(e1, e0))) | (~ spl24_102 | ~ spl24_125 | spl24_265)), inference(forward_demodulation, [], [f16099, f1119])).
fof(f16099, plain, (~ (e3 = op(e0, op(op(e0, e4), e0))) | (~ spl24_125 | spl24_265)), inference(forward_demodulation, [], [f1982, f1215])).
fof(f1982, plain, (~ (e3 = op(e0, op(op(e0, op(e0, e0)), e0))) | spl24_265), inference(avatar_component_clause, [], [f1980])).
fof(f1980, plain, (spl24_265 <=> (e3 = op(e0, op(op(e0, op(e0, e0)), e0)))), introduced(avatar_definition, [new_symbols(naming, [spl24_265])])).
fof(f16164, plain, (~ spl24_66 | ~ spl24_98 | spl24_192), inference(avatar_split_clause, [], [f16157, f1587, f1100, f966])).
fof(f966, plain, (spl24_66 <=> (e0 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_66])])).
fof(f16157, plain, (~ (e0 = op(e2, e1)) | (~ spl24_98 | spl24_192)), inference(backward_demodulation, [], [f1589, f1102])).
fof(f1589, plain, (~ (e0 = op(op(e1, e0), e1)) | spl24_192), inference(avatar_component_clause, [], [f1587])).
fof(f16163, plain, (~ spl24_78 | ~ spl24_98), inference(avatar_split_clause, [], [f16154, f1100, f1016])).
fof(f16154, plain, (~ (e2 = op(e1, e4)) | ~ spl24_98), inference(backward_demodulation, [], [f445, f1102])).
fof(f16152, plain, (~ spl24_78 | ~ spl24_12 | spl24_136), inference(avatar_split_clause, [], [f16151, f1312, f739, f1016])).
fof(f16151, plain, (~ (e2 = op(e1, e4)) | (~ spl24_12 | spl24_136)), inference(forward_demodulation, [], [f1314, f741])).
fof(f16147, plain, (~ spl24_66 | ~ spl24_86 | spl24_175), inference(avatar_split_clause, [], [f16146, f1503, f1050, f966])).
fof(f16146, plain, (~ (e0 = op(e2, e1)) | (~ spl24_86 | spl24_175)), inference(forward_demodulation, [], [f1505, f1052])).
fof(f16142, plain, (spl24_98 | ~ spl24_72 | ~ spl24_178), inference(avatar_split_clause, [], [f16141, f1517, f991, f1100])).
fof(f16141, plain, ((e2 = op(e1, e0)) | (~ spl24_72 | ~ spl24_178)), inference(forward_demodulation, [], [f1519, f993])).
fof(f16126, plain, (~ spl24_54 | ~ spl24_65 | spl24_222), inference(avatar_split_clause, [], [f16125, f1732, f961, f915])).
fof(f1732, plain, (spl24_222 <=> (e3 = op(e2, op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl24_222])])).
fof(f16125, plain, (~ (e3 = op(e2, e4)) | (~ spl24_65 | spl24_222)), inference(forward_demodulation, [], [f1734, f963])).
fof(f1734, plain, (~ (e3 = op(e2, op(e2, e2))) | spl24_222), inference(avatar_component_clause, [], [f1732])).
fof(f16118, plain, (spl24_79 | ~ spl24_95 | ~ spl24_232), inference(avatar_split_clause, [], [f16117, f1778, f1087, f1020])).
fof(f16117, plain, ((e3 = op(e1, e4)) | (~ spl24_95 | ~ spl24_232)), inference(forward_demodulation, [], [f1779, f1089])).
fof(f16116, plain, (~ spl24_78 | ~ spl24_95 | spl24_236), inference(avatar_split_clause, [], [f16115, f1800, f1087, f1016])).
fof(f16115, plain, (~ (e2 = op(e1, e4)) | (~ spl24_95 | spl24_236)), inference(forward_demodulation, [], [f1802, f1089])).
fof(f16107, plain, (spl24_51 | ~ spl24_65 | ~ spl24_253), inference(avatar_split_clause, [], [f16106, f1893, f961, f903])).
fof(f16106, plain, ((e0 = op(e2, e4)) | (~ spl24_65 | ~ spl24_253)), inference(forward_demodulation, [], [f1894, f963])).
fof(f1894, plain, ((e0 = op(e2, op(e2, e2))) | ~ spl24_253), inference(avatar_component_clause, [], [f1893])).
fof(f16089, plain, (~ spl24_67 | ~ spl24_42), inference(avatar_split_clause, [], [f16088, f865, f970])).
fof(f16088, plain, (~ (e1 = op(e2, e1)) | ~ spl24_42), inference(forward_demodulation, [], [f399, f867])).
fof(f867, plain, ((e1 = op(e3, e1)) | ~ spl24_42), inference(avatar_component_clause, [], [f865])).
fof(f16087, plain, (~ spl24_70 | ~ spl24_95), inference(avatar_split_clause, [], [f16086, f1087, f982])).
fof(f16086, plain, (~ (e4 = op(e2, e1)) | ~ spl24_95), inference(forward_demodulation, [], [f396, f1089])).
fof(f396, plain, ~ (op(e1, e1) = op(e2, e1)), inference(cnf_transformation, [], [f4])).
fof(f16085, plain, (~ spl24_68 | ~ spl24_118), inference(avatar_split_clause, [], [f16084, f1184, f974])).
fof(f16084, plain, (~ (e2 = op(e2, e1)) | ~ spl24_118), inference(forward_demodulation, [], [f393, f1186])).
fof(f16034, plain, (~ spl24_3 | ~ spl24_78), inference(avatar_split_clause, [], [f16029, f1016, f701])).
fof(f16029, plain, (~ (e2 = op(e4, e4)) | ~ spl24_78), inference(backward_demodulation, [], [f428, f1018])).
fof(f16027, plain, (~ spl24_11 | ~ spl24_86), inference(avatar_split_clause, [], [f16019, f1050, f735])).
fof(f16019, plain, (~ (e0 = op(e4, e2)) | ~ spl24_86), inference(backward_demodulation, [], [f408, f1052])).
fof(f16014, plain, (~ spl24_89 | ~ spl24_99), inference(avatar_split_clause, [], [f16009, f1104, f1062])).
fof(f16009, plain, (~ (e3 = op(e1, e2)) | ~ spl24_99), inference(backward_demodulation, [], [f443, f1106])).
fof(f15997, plain, (~ spl24_115 | ~ spl24_125), inference(avatar_split_clause, [], [f15994, f1213, f1171])).
fof(f15994, plain, (~ (e4 = op(e0, e2)) | ~ spl24_125), inference(backward_demodulation, [], [f433, f1215])).
fof(f15996, plain, (~ spl24_75 | ~ spl24_125), inference(avatar_split_clause, [], [f15992, f1213, f1003])).
fof(f15992, plain, (~ (e4 = op(e2, e0)) | ~ spl24_125), inference(backward_demodulation, [], [f383, f1215])).
fof(f15989, plain, (~ spl24_19 | ~ spl24_42 | ~ spl24_141), inference(avatar_contradiction_clause, [], [f15988])).
fof(f15988, plain, ($false | (~ spl24_19 | ~ spl24_42 | ~ spl24_141)), inference(subsumption_resolution, [], [f15987, f488])).
fof(f15987, plain, ((e1 = e4) | (~ spl24_19 | ~ spl24_42 | ~ spl24_141)), inference(forward_demodulation, [], [f15986, f867])).
fof(f15986, plain, ((e4 = op(e3, e1)) | (~ spl24_19 | ~ spl24_141)), inference(forward_demodulation, [], [f1338, f770])).
fof(f15954, plain, (spl24_78 | ~ spl24_95 | ~ spl24_236), inference(avatar_split_clause, [], [f15953, f1800, f1087, f1016])).
fof(f15953, plain, ((e2 = op(e1, e4)) | (~ spl24_95 | ~ spl24_236)), inference(forward_demodulation, [], [f1801, f1089])).
fof(f1801, plain, ((e2 = op(e1, op(e1, e1))) | ~ spl24_236), inference(avatar_component_clause, [], [f1800])).
fof(f15952, plain, (~ spl24_99 | ~ spl24_95 | ~ spl24_258 | spl24_264), inference(avatar_split_clause, [], [f15947, f1970, f1932, f1087, f1104])).
fof(f1932, plain, (spl24_258 <=> (e0 = op(op(e1, op(e1, e1)), e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_258])])).
fof(f1970, plain, (spl24_264 <=> (e3 = op(e1, op(op(e1, op(e1, e1)), e1)))), introduced(avatar_definition, [new_symbols(naming, [spl24_264])])).
fof(f15947, plain, (~ (e3 = op(e1, e0)) | (~ spl24_95 | ~ spl24_258 | spl24_264)), inference(backward_demodulation, [], [f15935, f15946])).
fof(f15946, plain, ((e0 = op(op(e1, e4), e1)) | (~ spl24_95 | ~ spl24_258)), inference(forward_demodulation, [], [f1933, f1089])).
fof(f1933, plain, ((e0 = op(op(e1, op(e1, e1)), e1)) | ~ spl24_258), inference(avatar_component_clause, [], [f1932])).
fof(f15935, plain, (~ (e3 = op(e1, op(op(e1, e4), e1))) | (~ spl24_95 | spl24_264)), inference(forward_demodulation, [], [f1972, f1089])).
fof(f1972, plain, (~ (e3 = op(e1, op(op(e1, op(e1, e1)), e1))) | spl24_264), inference(avatar_component_clause, [], [f1970])).
fof(f15944, plain, (~ spl24_98 | spl24_251 | ~ spl24_260), inference(avatar_split_clause, [], [f15939, f1944, f1882, f1100])).
fof(f15939, plain, (~ (e2 = op(e1, e0)) | (spl24_251 | ~ spl24_260)), inference(backward_demodulation, [], [f1884, f1945])).
fof(f1945, plain, ((e1 = op(e0, op(e0, e0))) | ~ spl24_260), inference(avatar_component_clause, [], [f1944])).
fof(f15937, plain, (~ spl24_76 | ~ spl24_95 | spl24_261), inference(avatar_split_clause, [], [f15936, f1949, f1087, f1008])).
fof(f15936, plain, (~ (e0 = op(e1, e4)) | (~ spl24_95 | spl24_261)), inference(forward_demodulation, [], [f1951, f1089])).
fof(f1951, plain, (~ (e0 = op(e1, op(e1, e1))) | spl24_261), inference(avatar_component_clause, [], [f1949])).
fof(f15882, plain, (~ spl24_3 | ~ spl24_11 | ~ spl24_102 | spl24_226), inference(avatar_contradiction_clause, [], [f15881])).
fof(f15881, plain, ($false | (~ spl24_3 | ~ spl24_11 | ~ spl24_102 | spl24_226)), inference(subsumption_resolution, [], [f15880, f1119])).
fof(f15880, plain, (~ (e1 = op(e0, e4)) | (~ spl24_3 | ~ spl24_11 | spl24_226)), inference(forward_demodulation, [], [f15870, f737])).
fof(f15870, plain, (~ (e1 = op(op(e4, e2), e4)) | (~ spl24_3 | spl24_226)), inference(backward_demodulation, [], [f1753, f703])).
fof(f1753, plain, (~ (e1 = op(op(e4, op(e4, e4)), e4)) | spl24_226), inference(avatar_component_clause, [], [f1751])).
fof(f1751, plain, (spl24_226 <=> (e1 = op(op(e4, op(e4, e4)), e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_226])])).
fof(f15841, plain, (~ spl24_10 | ~ spl24_30 | spl24_148), inference(avatar_split_clause, [], [f15835, f1369, f814, f730])).
fof(f15835, plain, (~ (e4 = op(e4, e3)) | (~ spl24_30 | spl24_148)), inference(backward_demodulation, [], [f1371, f816])).
fof(f816, plain, ((e4 = op(e3, e4)) | ~ spl24_30), inference(avatar_component_clause, [], [f814])).
fof(f15839, plain, (~ spl24_5 | ~ spl24_30), inference(avatar_split_clause, [], [f15833, f814, f709])).
fof(f15833, plain, (~ (e4 = op(e4, e4)) | ~ spl24_30), inference(backward_demodulation, [], [f431, f816])).
fof(f15823, plain, (~ spl24_29 | ~ spl24_34), inference(avatar_split_clause, [], [f15812, f831, f810])).
fof(f810, plain, (spl24_29 <=> (e3 = op(e3, e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_29])])).
fof(f15812, plain, (~ (e3 = op(e3, e4)) | ~ spl24_34), inference(backward_demodulation, [], [f471, f833])).
fof(f833, plain, ((e3 = op(e3, e3)) | ~ spl24_34), inference(avatar_component_clause, [], [f831])).
fof(f15800, plain, (~ spl24_29 | ~ spl24_54), inference(avatar_split_clause, [], [f15794, f915, f810])).
fof(f15794, plain, (~ (e3 = op(e3, e4)) | ~ spl24_54), inference(backward_demodulation, [], [f429, f917])).
fof(f15790, plain, (~ spl24_53 | ~ spl24_58), inference(avatar_split_clause, [], [f15782, f932, f911])).
fof(f911, plain, (spl24_53 <=> (e2 = op(e2, e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_53])])).
fof(f15782, plain, (~ (e2 = op(e2, e4)) | ~ spl24_58), inference(backward_demodulation, [], [f461, f934])).
fof(f934, plain, ((e2 = op(e2, e3)) | ~ spl24_58), inference(avatar_component_clause, [], [f932])).
fof(f461, plain, ~ (op(e2, e3) = op(e2, e4)), inference(cnf_transformation, [], [f4])).
fof(f15779, plain, (~ spl24_57 | ~ spl24_62), inference(avatar_split_clause, [], [f15773, f949, f928])).
fof(f15773, plain, (~ (e1 = op(e2, e3)) | ~ spl24_62), inference(backward_demodulation, [], [f459, f951])).
fof(f15746, plain, (~ spl24_77 | ~ spl24_82), inference(avatar_split_clause, [], [f15735, f1033, f1012])).
fof(f15735, plain, (~ (e1 = op(e1, e4)) | ~ spl24_82), inference(backward_demodulation, [], [f451, f1035])).
fof(f1035, plain, ((e1 = op(e1, e3)) | ~ spl24_82), inference(avatar_component_clause, [], [f1033])).
fof(f15731, plain, (~ spl24_43 | ~ spl24_89 | spl24_187), inference(avatar_split_clause, [], [f15728, f1561, f1062, f869])).
fof(f869, plain, (spl24_43 <=> (e2 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_43])])).
fof(f15728, plain, (~ (e2 = op(e3, e1)) | (~ spl24_89 | spl24_187)), inference(backward_demodulation, [], [f1563, f1064])).
fof(f1563, plain, (~ (e2 = op(op(e1, e2), e1)) | spl24_187), inference(avatar_component_clause, [], [f1561])).
fof(f15730, plain, (~ spl24_64 | ~ spl24_89), inference(avatar_split_clause, [], [f15722, f1062, f957])).
fof(f15722, plain, (~ (e3 = op(e2, e2)) | ~ spl24_89), inference(backward_demodulation, [], [f406, f1064])).
fof(f15721, plain, (~ spl24_85 | ~ spl24_95), inference(avatar_split_clause, [], [f15713, f1087, f1045])).
fof(f15713, plain, (~ (e4 = op(e1, e3)) | ~ spl24_95), inference(backward_demodulation, [], [f447, f1089])).
fof(f15690, plain, (~ spl24_101 | ~ spl24_106), inference(avatar_split_clause, [], [f15676, f1134, f1113])).
fof(f1113, plain, (spl24_101 <=> (e0 = op(e0, e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_101])])).
fof(f15676, plain, (~ (e0 = op(e0, e4)) | ~ spl24_106), inference(backward_demodulation, [], [f441, f1136])).
fof(f15671, plain, (~ spl24_54 | ~ spl24_115 | ~ spl24_253 | spl24_262), inference(avatar_split_clause, [], [f15667, f1960, f1893, f1171, f915])).
fof(f1960, plain, (spl24_262 <=> (e3 = op(e2, op(op(e2, op(e2, e2)), e2)))), introduced(avatar_definition, [new_symbols(naming, [spl24_262])])).
fof(f15667, plain, (~ (e3 = op(e2, e4)) | (~ spl24_115 | ~ spl24_253 | spl24_262)), inference(backward_demodulation, [], [f15575, f1173])).
fof(f15575, plain, (~ (e3 = op(e2, op(e0, e2))) | (~ spl24_253 | spl24_262)), inference(backward_demodulation, [], [f1962, f1894])).
fof(f1962, plain, (~ (e3 = op(e2, op(op(e2, op(e2, e2)), e2))) | spl24_262), inference(avatar_component_clause, [], [f1960])).
fof(f15670, plain, (spl24_11 | ~ spl24_115 | ~ spl24_203), inference(avatar_split_clause, [], [f15665, f1641, f1171, f735])).
fof(f1641, plain, (spl24_203 <=> (e0 = op(op(e0, e2), e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_203])])).
fof(f15665, plain, ((e0 = op(e4, e2)) | (~ spl24_115 | ~ spl24_203)), inference(backward_demodulation, [], [f1643, f1173])).
fof(f1643, plain, ((e0 = op(op(e0, e2), e2)) | ~ spl24_203), inference(avatar_component_clause, [], [f1641])).
fof(f15669, plain, (~ spl24_40 | ~ spl24_115), inference(avatar_split_clause, [], [f15660, f1171, f856])).
fof(f15660, plain, (~ (e4 = op(e3, e2)) | ~ spl24_115), inference(backward_demodulation, [], [f404, f1173])).
fof(f404, plain, ~ (op(e0, e2) = op(e3, e2)), inference(cnf_transformation, [], [f4])).
fof(f15657, plain, (~ spl24_108 | ~ spl24_118), inference(avatar_split_clause, [], [f15650, f1184, f1142])).
fof(f15650, plain, (~ (e2 = op(e0, e3)) | ~ spl24_118), inference(backward_demodulation, [], [f437, f1186])).
fof(f15656, plain, (~ spl24_43 | ~ spl24_118), inference(avatar_split_clause, [], [f15647, f1184, f869])).
fof(f15647, plain, (~ (e2 = op(e3, e1)) | ~ spl24_118), inference(backward_demodulation, [], [f394, f1186])).
fof(f394, plain, ~ (op(e0, e1) = op(e3, e1)), inference(cnf_transformation, [], [f4])).
fof(f15601, plain, (~ spl24_112 | ~ spl24_66 | spl24_173), inference(avatar_split_clause, [], [f15600, f1493, f966, f1159])).
fof(f1493, plain, (spl24_173 <=> (e1 = op(op(e2, e1), e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_173])])).
fof(f15600, plain, (~ (e1 = op(e0, e2)) | (~ spl24_66 | spl24_173)), inference(forward_demodulation, [], [f1495, f968])).
fof(f968, plain, ((e0 = op(e2, e1)) | ~ spl24_66), inference(avatar_component_clause, [], [f966])).
fof(f1495, plain, (~ (e1 = op(op(e2, e1), e2)) | spl24_173), inference(avatar_component_clause, [], [f1493])).
fof(f15599, plain, (spl24_118 | ~ spl24_66 | ~ spl24_174), inference(avatar_split_clause, [], [f15598, f1498, f966, f1184])).
fof(f1498, plain, (spl24_174 <=> (e2 = op(op(e2, e1), e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_174])])).
fof(f15598, plain, ((e2 = op(e0, e1)) | (~ spl24_66 | ~ spl24_174)), inference(forward_demodulation, [], [f1500, f968])).
fof(f1500, plain, ((e2 = op(op(e2, e1), e1)) | ~ spl24_174), inference(avatar_component_clause, [], [f1498])).
fof(f15593, plain, (spl24_118 | ~ spl24_98 | ~ spl24_194), inference(avatar_split_clause, [], [f15592, f1597, f1100, f1184])).
fof(f15592, plain, ((e2 = op(e0, e1)) | (~ spl24_98 | ~ spl24_194)), inference(forward_demodulation, [], [f1598, f1102])).
fof(f15581, plain, (~ spl24_108 | ~ spl24_124 | spl24_252), inference(avatar_split_clause, [], [f15580, f1888, f1209, f1142])).
fof(f15580, plain, (~ (e2 = op(e0, e3)) | (~ spl24_124 | spl24_252)), inference(forward_demodulation, [], [f1890, f1211])).
fof(f1890, plain, (~ (e2 = op(e0, op(e0, e0))) | spl24_252), inference(avatar_component_clause, [], [f1888])).
fof(f15579, plain, (~ spl24_112 | spl24_234 | ~ spl24_253), inference(avatar_split_clause, [], [f15574, f1893, f1788, f1159])).
fof(f1788, plain, (spl24_234 <=> (e1 = op(op(e2, op(e2, e2)), e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_234])])).
fof(f15574, plain, (~ (e1 = op(e0, e2)) | (spl24_234 | ~ spl24_253)), inference(backward_demodulation, [], [f1790, f1894])).
fof(f1790, plain, (~ (e1 = op(op(e2, op(e2, e2)), e2)) | spl24_234), inference(avatar_component_clause, [], [f1788])).
fof(f15513, plain, (~ spl24_32 | ~ spl24_47), inference(avatar_split_clause, [], [f15506, f886, f823])).
fof(f15506, plain, (~ (e1 = op(e3, e3)) | ~ spl24_47), inference(backward_demodulation, [], [f464, f888])).
fof(f464, plain, ~ (op(e3, e0) = op(e3, e3)), inference(cnf_transformation, [], [f4])).
fof(f15475, plain, (~ spl24_56 | ~ spl24_66), inference(avatar_split_clause, [], [f15467, f966, f924])).
fof(f15467, plain, (~ (e0 = op(e2, e3)) | ~ spl24_66), inference(backward_demodulation, [], [f457, f968])).
fof(f15462, plain, (~ spl24_11 | ~ spl24_75 | spl24_177), inference(avatar_split_clause, [], [f15456, f1512, f1003, f735])).
fof(f1512, plain, (spl24_177 <=> (e0 = op(op(e2, e0), e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_177])])).
fof(f15456, plain, (~ (e0 = op(e4, e2)) | (~ spl24_75 | spl24_177)), inference(backward_demodulation, [], [f1514, f1005])).
fof(f1514, plain, (~ (e0 = op(op(e2, e0), e2)) | spl24_177), inference(avatar_component_clause, [], [f1512])).
fof(f15461, plain, (~ spl24_70 | ~ spl24_75), inference(avatar_split_clause, [], [f15452, f1003, f982])).
fof(f15452, plain, (~ (e4 = op(e2, e1)) | ~ spl24_75), inference(backward_demodulation, [], [f452, f1005])).
fof(f15420, plain, (~ spl24_89 | ~ spl24_94), inference(avatar_split_clause, [], [f15412, f1083, f1062])).
fof(f15412, plain, (~ (e3 = op(e1, e2)) | ~ spl24_94), inference(backward_demodulation, [], [f446, f1085])).
fof(f446, plain, ~ (op(e1, e1) = op(e1, e2)), inference(cnf_transformation, [], [f4])).
fof(f15407, plain, (~ spl24_73 | ~ spl24_98), inference(avatar_split_clause, [], [f15398, f1100, f995])).
fof(f15398, plain, (~ (e2 = op(e2, e0)) | ~ spl24_98), inference(backward_demodulation, [], [f386, f1102])).
fof(f386, plain, ~ (op(e1, e0) = op(e2, e0)), inference(cnf_transformation, [], [f4])).
fof(f15376, plain, (~ spl24_56 | ~ spl24_108 | spl24_200), inference(avatar_split_clause, [], [f15374, f1626, f1142, f924])).
fof(f15374, plain, (~ (e0 = op(e2, e3)) | (~ spl24_108 | spl24_200)), inference(backward_demodulation, [], [f1627, f1144])).
fof(f15363, plain, (~ spl24_62 | ~ spl24_112), inference(avatar_split_clause, [], [f15353, f1159, f949])).
fof(f15353, plain, (~ (e1 = op(e2, e2)) | ~ spl24_112), inference(backward_demodulation, [], [f403, f1161])).
fof(f15351, plain, (~ spl24_105 | ~ spl24_120), inference(avatar_split_clause, [], [f15346, f1192, f1129])).
fof(f15346, plain, (~ (e4 = op(e0, e4)) | ~ spl24_120), inference(backward_demodulation, [], [f438, f1194])).
fof(f438, plain, ~ (op(e0, e1) = op(e0, e4)), inference(cnf_transformation, [], [f4])).
fof(f15350, plain, (~ spl24_70 | ~ spl24_120), inference(avatar_split_clause, [], [f15342, f1192, f982])).
fof(f15342, plain, (~ (e4 = op(e2, e1)) | ~ spl24_120), inference(backward_demodulation, [], [f393, f1194])).
fof(f15339, plain, (~ spl24_109 | ~ spl24_124), inference(avatar_split_clause, [], [f15333, f1209, f1146])).
fof(f15333, plain, (~ (e3 = op(e0, e3)) | ~ spl24_124), inference(backward_demodulation, [], [f434, f1211])).
fof(f15326, plain, (spl24_5 | ~ spl24_130), inference(avatar_split_clause, [], [f15316, f1234, f709])).
fof(f1234, plain, (spl24_130 <=> (e4 = unit)), introduced(avatar_definition, [new_symbols(naming, [spl24_130])])).
fof(f15316, plain, ((e4 = op(e4, e4)) | ~ spl24_130), inference(backward_demodulation, [], [f330, f1236])).
fof(f1236, plain, ((e4 = unit) | ~ spl24_130), inference(avatar_component_clause, [], [f1234])).
fof(f15324, plain, (spl24_29 | ~ spl24_130), inference(avatar_split_clause, [], [f15314, f1234, f810])).
fof(f15314, plain, ((e3 = op(e3, e4)) | ~ spl24_130), inference(backward_demodulation, [], [f328, f1236])).
fof(f15323, plain, (spl24_9 | ~ spl24_130), inference(avatar_split_clause, [], [f15313, f1234, f726])).
fof(f15313, plain, ((e3 = op(e4, e3)) | ~ spl24_130), inference(backward_demodulation, [], [f327, f1236])).
fof(f15322, plain, (spl24_53 | ~ spl24_130), inference(avatar_split_clause, [], [f15312, f1234, f911])).
fof(f15312, plain, ((e2 = op(e2, e4)) | ~ spl24_130), inference(backward_demodulation, [], [f326, f1236])).
fof(f15321, plain, (spl24_13 | ~ spl24_130), inference(avatar_split_clause, [], [f15311, f1234, f743])).
fof(f15311, plain, ((e2 = op(e4, e2)) | ~ spl24_130), inference(backward_demodulation, [], [f325, f1236])).
fof(f15320, plain, (spl24_77 | ~ spl24_130), inference(avatar_split_clause, [], [f15310, f1234, f1012])).
fof(f15310, plain, ((e1 = op(e1, e4)) | ~ spl24_130), inference(backward_demodulation, [], [f324, f1236])).
fof(f15319, plain, (spl24_17 | ~ spl24_130), inference(avatar_split_clause, [], [f15309, f1234, f760])).
fof(f760, plain, (spl24_17 <=> (e1 = op(e4, e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_17])])).
fof(f15309, plain, ((e1 = op(e4, e1)) | ~ spl24_130), inference(backward_demodulation, [], [f323, f1236])).
fof(f15318, plain, (spl24_101 | ~ spl24_130), inference(avatar_split_clause, [], [f15308, f1234, f1113])).
fof(f15308, plain, ((e0 = op(e0, e4)) | ~ spl24_130), inference(backward_demodulation, [], [f322, f1236])).
fof(f15317, plain, (spl24_21 | ~ spl24_130), inference(avatar_split_clause, [], [f15307, f1234, f777])).
fof(f777, plain, (spl24_21 <=> (e0 = op(e4, e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_21])])).
fof(f15307, plain, ((e0 = op(e4, e0)) | ~ spl24_130), inference(backward_demodulation, [], [f321, f1236])).
fof(f15285, plain, (spl24_57 | ~ spl24_43 | ~ spl24_157), inference(avatar_split_clause, [], [f15284, f1414, f869, f928])).
fof(f1414, plain, (spl24_157 <=> (e1 = op(op(e3, e1), e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_157])])).
fof(f15284, plain, ((e1 = op(e2, e3)) | (~ spl24_43 | ~ spl24_157)), inference(forward_demodulation, [], [f1415, f871])).
fof(f871, plain, ((e2 = op(e3, e1)) | ~ spl24_43), inference(avatar_component_clause, [], [f869])).
fof(f1415, plain, ((e1 = op(op(e3, e1), e3)) | ~ spl24_157), inference(avatar_component_clause, [], [f1414])).
fof(f15283, plain, (~ spl24_19 | ~ spl24_85 | spl24_184), inference(avatar_split_clause, [], [f15282, f1546, f1045, f768])).
fof(f15282, plain, (~ (e3 = op(e4, e1)) | (~ spl24_85 | spl24_184)), inference(forward_demodulation, [], [f1548, f1047])).
fof(f15276, plain, (~ spl24_19 | spl24_228 | ~ spl24_229), inference(avatar_split_clause, [], [f15272, f1764, f1760, f768])).
fof(f1760, plain, (spl24_228 <=> (e3 = op(op(e1, op(e1, e1)), e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_228])])).
fof(f1764, plain, (spl24_229 <=> (e4 = op(e1, op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl24_229])])).
fof(f15272, plain, (~ (e3 = op(e4, e1)) | (spl24_228 | ~ spl24_229)), inference(backward_demodulation, [], [f1762, f1765])).
fof(f1765, plain, ((e4 = op(e1, op(e1, e1))) | ~ spl24_229), inference(avatar_component_clause, [], [f1764])).
fof(f1762, plain, (~ (e3 = op(op(e1, op(e1, e1)), e1)) | spl24_228), inference(avatar_component_clause, [], [f1760])).
fof(f15205, plain, (~ spl24_26 | spl24_104 | ~ spl24_149), inference(avatar_contradiction_clause, [], [f15204])).
fof(f15204, plain, ($false | (~ spl24_26 | spl24_104 | ~ spl24_149)), inference(subsumption_resolution, [], [f15201, f1126])).
fof(f1126, plain, (~ (e3 = op(e0, e4)) | spl24_104), inference(avatar_component_clause, [], [f1125])).
fof(f15201, plain, ((e3 = op(e0, e4)) | (~ spl24_26 | ~ spl24_149)), inference(backward_demodulation, [], [f1376, f800])).
fof(f1376, plain, ((e3 = op(op(e3, e4), e4)) | ~ spl24_149), inference(avatar_component_clause, [], [f1374])).
fof(f1374, plain, (spl24_149 <=> (e3 = op(op(e3, e4), e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_149])])).
fof(f15194, plain, (~ spl24_27 | ~ spl24_32), inference(avatar_split_clause, [], [f15187, f823, f802])).
fof(f15187, plain, (~ (e1 = op(e3, e4)) | ~ spl24_32), inference(backward_demodulation, [], [f471, f825])).
fof(f825, plain, ((e1 = op(e3, e3)) | ~ spl24_32), inference(avatar_component_clause, [], [f823])).
fof(f15160, plain, (~ spl24_11 | ~ spl24_62 | ~ spl24_70 | spl24_250), inference(avatar_split_clause, [], [f15159, f1876, f982, f949, f735])).
fof(f1876, plain, (spl24_250 <=> (e0 = op(op(e2, op(e2, e2)), e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_250])])).
fof(f15159, plain, (~ (e0 = op(e4, e2)) | (~ spl24_62 | ~ spl24_70 | spl24_250)), inference(forward_demodulation, [], [f15152, f984])).
fof(f15152, plain, (~ (e0 = op(op(e2, e1), e2)) | (~ spl24_62 | spl24_250)), inference(backward_demodulation, [], [f1878, f951])).
fof(f1878, plain, (~ (e0 = op(op(e2, op(e2, e2)), e2)) | spl24_250), inference(avatar_component_clause, [], [f1876])).
fof(f15156, plain, (~ spl24_62 | ~ spl24_70 | spl24_218), inference(avatar_contradiction_clause, [], [f15155])).
fof(f15155, plain, ($false | (~ spl24_62 | ~ spl24_70 | spl24_218)), inference(subsumption_resolution, [], [f15147, f984])).
fof(f15147, plain, (~ (e4 = op(e2, e1)) | (~ spl24_62 | spl24_218)), inference(backward_demodulation, [], [f1716, f951])).
fof(f1716, plain, (~ (e4 = op(e2, op(e2, e2))) | spl24_218), inference(avatar_component_clause, [], [f1714])).
fof(f1714, plain, (spl24_218 <=> (e4 = op(e2, op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl24_218])])).
fof(f15060, plain, (~ spl24_107 | ~ spl24_117), inference(avatar_split_clause, [], [f15047, f1180, f1138])).
fof(f15047, plain, (~ (e1 = op(e0, e3)) | ~ spl24_117), inference(backward_demodulation, [], [f437, f1182])).
fof(f14973, plain, (spl24_54 | ~ spl24_8 | ~ spl24_132), inference(avatar_split_clause, [], [f14972, f1293, f722, f915])).
fof(f14972, plain, ((e3 = op(e2, e4)) | (~ spl24_8 | ~ spl24_132)), inference(forward_demodulation, [], [f1294, f724])).
fof(f1294, plain, ((e3 = op(op(e4, e3), e4)) | ~ spl24_132), inference(avatar_component_clause, [], [f1293])).
fof(f14971, plain, (~ spl24_27 | ~ spl24_19 | spl24_140), inference(avatar_split_clause, [], [f14970, f1331, f768, f802])).
fof(f1331, plain, (spl24_140 <=> (e1 = op(op(e4, e1), e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_140])])).
fof(f14970, plain, (~ (e1 = op(e3, e4)) | (~ spl24_19 | spl24_140)), inference(forward_demodulation, [], [f1333, f770])).
fof(f1333, plain, (~ (e1 = op(op(e4, e1), e4)) | spl24_140), inference(avatar_component_clause, [], [f1331])).
fof(f14927, plain, (~ spl24_50 | ~ spl24_241 | spl24_267), inference(avatar_split_clause, [], [f14925, f2005, f1834, f898])).
fof(f14925, plain, (~ (e4 = op(e3, e0)) | (~ spl24_241 | spl24_267)), inference(backward_demodulation, [], [f2007, f1835])).
fof(f14900, plain, (~ spl24_68 | ~ spl24_43), inference(avatar_split_clause, [], [f14899, f869, f974])).
fof(f14899, plain, (~ (e2 = op(e2, e1)) | ~ spl24_43), inference(forward_demodulation, [], [f399, f871])).
fof(f14873, plain, (~ spl24_2 | ~ spl24_22), inference(avatar_split_clause, [], [f14867, f781, f697])).
fof(f14867, plain, (~ (e1 = op(e4, e4)) | ~ spl24_22), inference(backward_demodulation, [], [f475, f783])).
fof(f783, plain, ((e1 = op(e4, e0)) | ~ spl24_22), inference(avatar_component_clause, [], [f781])).
fof(f14844, plain, (~ spl24_33 | ~ spl24_43), inference(avatar_split_clause, [], [f14838, f869, f827])).
fof(f14838, plain, (~ (e2 = op(e3, e3)) | ~ spl24_43), inference(backward_demodulation, [], [f467, f871])).
fof(f14837, plain, (~ spl24_6 | ~ spl24_50 | spl24_161), inference(avatar_split_clause, [], [f14834, f1433, f898, f714])).
fof(f1433, plain, (spl24_161 <=> (e0 = op(op(e3, e0), e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_161])])).
fof(f14834, plain, (~ (e0 = op(e4, e3)) | (~ spl24_50 | spl24_161)), inference(backward_demodulation, [], [f1435, f900])).
fof(f1435, plain, (~ (e0 = op(op(e3, e0), e3)) | spl24_161), inference(avatar_component_clause, [], [f1433])).
fof(f14827, plain, (spl24_27 | ~ spl24_79 | ~ spl24_182), inference(avatar_split_clause, [], [f14821, f1536, f1020, f802])).
fof(f14821, plain, ((e1 = op(e3, e4)) | (~ spl24_79 | ~ spl24_182)), inference(backward_demodulation, [], [f1538, f1022])).
fof(f14826, plain, (~ spl24_45 | ~ spl24_79 | spl24_181), inference(avatar_split_clause, [], [f14820, f1531, f1020, f877])).
fof(f14820, plain, (~ (e4 = op(e3, e1)) | (~ spl24_79 | spl24_181)), inference(backward_demodulation, [], [f1533, f1022])).
fof(f14784, plain, (~ spl24_23 | ~ spl24_103 | spl24_146), inference(avatar_split_clause, [], [f14780, f1360, f1121, f785])).
fof(f14780, plain, (~ (e2 = op(e4, e0)) | (~ spl24_103 | spl24_146)), inference(backward_demodulation, [], [f1362, f1123])).
fof(f14783, plain, (~ spl24_78 | ~ spl24_103), inference(avatar_split_clause, [], [f14775, f1121, f1016])).
fof(f14775, plain, (~ (e2 = op(e1, e4)) | ~ spl24_103), inference(backward_demodulation, [], [f422, f1123])).
fof(f14773, plain, (~ spl24_45 | ~ spl24_120), inference(avatar_split_clause, [], [f14768, f1192, f877])).
fof(f14768, plain, (~ (e4 = op(e3, e1)) | ~ spl24_120), inference(backward_demodulation, [], [f394, f1194])).
fof(f14743, plain, (~ spl24_99 | ~ spl24_107 | spl24_199), inference(avatar_split_clause, [], [f14742, f1621, f1138, f1104])).
fof(f14742, plain, (~ (e3 = op(e1, e0)) | (~ spl24_107 | spl24_199)), inference(forward_demodulation, [], [f1623, f1140])).
fof(f14726, plain, (~ spl24_99 | ~ spl24_91 | spl24_232), inference(avatar_split_clause, [], [f14725, f1778, f1071, f1104])).
fof(f14725, plain, (~ (e3 = op(e1, e0)) | (~ spl24_91 | spl24_232)), inference(forward_demodulation, [], [f1780, f1073])).
fof(f14724, plain, (spl24_98 | ~ spl24_91 | ~ spl24_236), inference(avatar_split_clause, [], [f14723, f1800, f1071, f1100])).
fof(f14723, plain, ((e2 = op(e1, e0)) | (~ spl24_91 | ~ spl24_236)), inference(forward_demodulation, [], [f1801, f1073])).
fof(f14717, plain, (~ spl24_99 | spl24_244 | ~ spl24_260), inference(avatar_split_clause, [], [f14711, f1944, f1848, f1104])).
fof(f14711, plain, (~ (e3 = op(e1, e0)) | (spl24_244 | ~ spl24_260)), inference(backward_demodulation, [], [f1850, f1945])).
fof(f14644, plain, (~ spl24_23 | ~ spl24_26 | ~ spl24_215 | spl24_255), inference(avatar_split_clause, [], [f14638, f1909, f1701, f798, f785])).
fof(f1701, plain, (spl24_215 <=> (e3 = op(e4, op(e4, e4)))), introduced(avatar_definition, [new_symbols(naming, [spl24_215])])).
fof(f1909, plain, (spl24_255 <=> (e2 = op(e4, op(op(e4, op(e4, e4)), e4)))), introduced(avatar_definition, [new_symbols(naming, [spl24_255])])).
fof(f14638, plain, (~ (e2 = op(e4, e0)) | (~ spl24_26 | ~ spl24_215 | spl24_255)), inference(backward_demodulation, [], [f14360, f800])).
fof(f14360, plain, (~ (e2 = op(e4, op(e3, e4))) | (~ spl24_215 | spl24_255)), inference(backward_demodulation, [], [f1911, f1702])).
fof(f1702, plain, ((e3 = op(e4, op(e4, e4))) | ~ spl24_215), inference(avatar_component_clause, [], [f1701])).
fof(f1911, plain, (~ (e2 = op(e4, op(op(e4, op(e4, e4)), e4))) | spl24_255), inference(avatar_component_clause, [], [f1909])).
fof(f14602, plain, (~ spl24_5 | ~ spl24_55), inference(avatar_split_clause, [], [f14595, f919, f709])).
fof(f14595, plain, (~ (e4 = op(e4, e4)) | ~ spl24_55), inference(backward_demodulation, [], [f430, f921])).
fof(f14592, plain, (spl24_33 | ~ spl24_59 | ~ spl24_169), inference(avatar_split_clause, [], [f14587, f1472, f936, f827])).
fof(f14587, plain, ((e2 = op(e3, e3)) | (~ spl24_59 | ~ spl24_169)), inference(backward_demodulation, [], [f1474, f938])).
fof(f14591, plain, (~ spl24_39 | ~ spl24_59 | spl24_168), inference(avatar_split_clause, [], [f14586, f1467, f936, f852])).
fof(f14586, plain, (~ (e3 = op(e3, e2)) | (~ spl24_59 | spl24_168)), inference(backward_demodulation, [], [f1469, f938])).
fof(f1469, plain, (~ (e3 = op(op(e2, e3), e2)) | spl24_168), inference(avatar_component_clause, [], [f1467])).
fof(f14589, plain, (~ spl24_9 | ~ spl24_59), inference(avatar_split_clause, [], [f14582, f936, f726])).
fof(f14582, plain, (~ (e3 = op(e4, e3)) | ~ spl24_59), inference(backward_demodulation, [], [f420, f938])).
fof(f420, plain, ~ (op(e2, e3) = op(e4, e3)), inference(cnf_transformation, [], [f4])).
fof(f14562, plain, (~ spl24_70 | ~ spl24_78 | spl24_181), inference(avatar_split_clause, [], [f14559, f1531, f1016, f982])).
fof(f14559, plain, (~ (e4 = op(e2, e1)) | (~ spl24_78 | spl24_181)), inference(backward_demodulation, [], [f1533, f1018])).
fof(f14561, plain, (~ spl24_53 | ~ spl24_78), inference(avatar_split_clause, [], [f14553, f1016, f911])).
fof(f14553, plain, (~ (e2 = op(e2, e4)) | ~ spl24_78), inference(backward_demodulation, [], [f426, f1018])).
fof(f14532, plain, (~ spl24_77 | ~ spl24_87), inference(avatar_split_clause, [], [f14522, f1054, f1012])).
fof(f14522, plain, (~ (e1 = op(e1, e4)) | ~ spl24_87), inference(backward_demodulation, [], [f450, f1056])).
fof(f14517, plain, (~ spl24_86 | ~ spl24_91), inference(avatar_split_clause, [], [f14507, f1071, f1050])).
fof(f14507, plain, (~ (e0 = op(e1, e2)) | ~ spl24_91), inference(backward_demodulation, [], [f446, f1073])).
fof(f14516, plain, (~ spl24_41 | ~ spl24_91), inference(avatar_split_clause, [], [f14505, f1071, f861])).
fof(f14505, plain, (~ (e0 = op(e3, e1)) | ~ spl24_91), inference(backward_demodulation, [], [f397, f1073])).
fof(f397, plain, ~ (op(e1, e1) = op(e3, e1)), inference(cnf_transformation, [], [f4])).
fof(f14503, plain, (~ spl24_47 | ~ spl24_99 | spl24_193), inference(avatar_split_clause, [], [f14500, f1592, f1104, f886])).
fof(f14500, plain, (~ (e1 = op(e3, e0)) | (~ spl24_99 | spl24_193)), inference(backward_demodulation, [], [f1593, f1106])).
fof(f14502, plain, (~ spl24_41 | ~ spl24_99 | spl24_192), inference(avatar_split_clause, [], [f14499, f1587, f1104, f861])).
fof(f14499, plain, (~ (e0 = op(e3, e1)) | (~ spl24_99 | spl24_192)), inference(backward_demodulation, [], [f1589, f1106])).
fof(f14501, plain, (~ spl24_94 | ~ spl24_99), inference(avatar_split_clause, [], [f14494, f1104, f1083])).
fof(f14494, plain, (~ (e3 = op(e1, e1)) | ~ spl24_99), inference(backward_demodulation, [], [f442, f1106])).
fof(f14492, plain, (spl24_26 | ~ spl24_104 | ~ spl24_197), inference(avatar_split_clause, [], [f14490, f1611, f1125, f798])).
fof(f14490, plain, ((e0 = op(e3, e4)) | (~ spl24_104 | ~ spl24_197)), inference(backward_demodulation, [], [f1613, f1127])).
fof(f14491, plain, (~ spl24_29 | ~ spl24_104), inference(avatar_split_clause, [], [f14485, f1125, f810])).
fof(f14485, plain, (~ (e3 = op(e3, e4)) | ~ spl24_104), inference(backward_demodulation, [], [f424, f1127])).
fof(f14481, plain, (~ spl24_107 | ~ spl24_118 | ~ spl24_163 | ~ spl24_248 | spl24_257), inference(avatar_contradiction_clause, [], [f14480])).
fof(f14480, plain, ($false | (~ spl24_107 | ~ spl24_118 | ~ spl24_163 | ~ spl24_248 | spl24_257)), inference(subsumption_resolution, [], [f14474, f1186])).
fof(f14474, plain, (~ (e2 = op(e0, e1)) | (~ spl24_107 | ~ spl24_163 | ~ spl24_248 | spl24_257)), inference(backward_demodulation, [], [f14383, f1140])).
fof(f14383, plain, (~ (e2 = op(e0, op(e0, e3))) | (~ spl24_163 | ~ spl24_248 | spl24_257)), inference(backward_demodulation, [], [f14337, f1444])).
fof(f14337, plain, (~ (e2 = op(e0, op(e3, e0))) | (~ spl24_248 | spl24_257)), inference(backward_demodulation, [], [f1926, f1867])).
fof(f1926, plain, (~ (e2 = op(e0, op(op(e0, op(e0, e0)), e0))) | spl24_257), inference(avatar_component_clause, [], [f1924])).
fof(f1924, plain, (spl24_257 <=> (e2 = op(e0, op(op(e0, op(e0, e0)), e0)))), introduced(avatar_definition, [new_symbols(naming, [spl24_257])])).
fof(f14477, plain, (~ spl24_32 | ~ spl24_107), inference(avatar_split_clause, [], [f14461, f1138, f823])).
fof(f14461, plain, (~ (e1 = op(e3, e3)) | ~ spl24_107), inference(backward_demodulation, [], [f414, f1140])).
fof(f14454, plain, (~ spl24_101 | ~ spl24_111), inference(avatar_split_clause, [], [f14437, f1155, f1113])).
fof(f14437, plain, (~ (e0 = op(e0, e4)) | ~ spl24_111), inference(backward_demodulation, [], [f440, f1157])).
fof(f440, plain, ~ (op(e0, e2) = op(e0, e4)), inference(cnf_transformation, [], [f4])).
fof(f14363, plain, (~ spl24_26 | ~ spl24_215 | spl24_242), inference(avatar_split_clause, [], [f14359, f1839, f1701, f798])).
fof(f1839, plain, (spl24_242 <=> (e0 = op(op(e4, op(e4, e4)), e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_242])])).
fof(f14359, plain, (~ (e0 = op(e3, e4)) | (~ spl24_215 | spl24_242)), inference(backward_demodulation, [], [f1841, f1702])).
fof(f1841, plain, (~ (e0 = op(op(e4, op(e4, e4)), e4)) | spl24_242), inference(avatar_component_clause, [], [f1839])).
fof(f14328, plain, (~ spl24_113 | ~ spl24_118), inference(avatar_split_clause, [], [f14327, f1184, f1163])).
fof(f14327, plain, (~ (e2 = op(e0, e2)) | ~ spl24_118), inference(forward_demodulation, [], [f436, f1186])).
fof(f14258, plain, (~ spl24_52 | ~ spl24_72), inference(avatar_split_clause, [], [f14252, f991, f907])).
fof(f14252, plain, (~ (e1 = op(e2, e4)) | ~ spl24_72), inference(backward_demodulation, [], [f455, f993])).
fof(f455, plain, ~ (op(e2, e0) = op(e2, e4)), inference(cnf_transformation, [], [f4])).
fof(f14235, plain, (~ spl24_68 | ~ spl24_83 | ~ spl24_159), inference(avatar_split_clause, [], [f14228, f1424, f1037, f974])).
fof(f14228, plain, (~ (e2 = op(e2, e1)) | (~ spl24_83 | ~ spl24_159)), inference(backward_demodulation, [], [f14083, f1039])).
fof(f14083, plain, (~ (op(e1, e3) = op(e2, e1)) | ~ spl24_159), inference(forward_demodulation, [], [f399, f1425])).
fof(f1425, plain, ((op(e1, e3) = op(e3, e1)) | ~ spl24_159), inference(avatar_component_clause, [], [f1424])).
fof(f14229, plain, (~ spl24_8 | ~ spl24_83), inference(avatar_split_clause, [], [f14219, f1037, f722])).
fof(f14219, plain, (~ (e2 = op(e4, e3)) | ~ spl24_83), inference(backward_demodulation, [], [f418, f1039])).
fof(f14218, plain, (~ spl24_86 | ~ spl24_118 | spl24_187), inference(avatar_contradiction_clause, [], [f14217])).
fof(f14217, plain, ($false | (~ spl24_86 | ~ spl24_118 | spl24_187)), inference(subsumption_resolution, [], [f14214, f1186])).
fof(f14214, plain, (~ (e2 = op(e0, e1)) | (~ spl24_86 | spl24_187)), inference(backward_demodulation, [], [f1563, f1052])).
fof(f14205, plain, (spl24_83 | ~ spl24_94 | ~ spl24_236), inference(avatar_split_clause, [], [f14201, f1800, f1083, f1037])).
fof(f14201, plain, ((e2 = op(e1, e3)) | (~ spl24_94 | ~ spl24_236)), inference(backward_demodulation, [], [f1801, f1085])).
fof(f14171, plain, (~ spl24_83 | ~ spl24_118 | ~ spl24_194), inference(avatar_split_clause, [], [f14163, f1597, f1184, f1037])).
fof(f14163, plain, (~ (e2 = op(e1, e3)) | (~ spl24_118 | ~ spl24_194)), inference(backward_demodulation, [], [f13452, f1186])).
fof(f13452, plain, (~ (op(e0, e1) = op(e1, e3)) | ~ spl24_194), inference(backward_demodulation, [], [f444, f1598])).
fof(f14166, plain, (~ spl24_103 | ~ spl24_118), inference(avatar_split_clause, [], [f14154, f1184, f1121])).
fof(f14154, plain, (~ (e2 = op(e0, e4)) | ~ spl24_118), inference(backward_demodulation, [], [f438, f1186])).
fof(f14107, plain, (spl24_72 | ~ spl24_112 | ~ spl24_179), inference(avatar_split_clause, [], [f14106, f1522, f1159, f991])).
fof(f14106, plain, ((e1 = op(e2, e0)) | (~ spl24_112 | ~ spl24_179)), inference(forward_demodulation, [], [f1523, f1161])).
fof(f1523, plain, ((op(e0, e2) = op(e2, e0)) | ~ spl24_179), inference(avatar_component_clause, [], [f1522])).
fof(f14104, plain, (spl24_86 | ~ spl24_112 | ~ spl24_203), inference(avatar_split_clause, [], [f14103, f1641, f1159, f1050])).
fof(f14103, plain, ((e0 = op(e1, e2)) | (~ spl24_112 | ~ spl24_203)), inference(forward_demodulation, [], [f1643, f1161])).
fof(f14101, plain, (~ spl24_70 | spl24_231 | ~ spl24_236), inference(avatar_split_clause, [], [f14095, f1800, f1774, f982])).
fof(f1774, plain, (spl24_231 <=> (e4 = op(op(e1, op(e1, e1)), e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_231])])).
fof(f14095, plain, (~ (e4 = op(e2, e1)) | (spl24_231 | ~ spl24_236)), inference(backward_demodulation, [], [f1776, f1801])).
fof(f1776, plain, (~ (e4 = op(op(e1, op(e1, e1)), e1)) | spl24_231), inference(avatar_component_clause, [], [f1774])).
fof(f14092, plain, (~ spl24_22 | ~ spl24_110 | ~ spl24_124 | spl24_259), inference(avatar_split_clause, [], [f14091, f1938, f1209, f1150, f781])).
fof(f14091, plain, (~ (e1 = op(e4, e0)) | (~ spl24_110 | ~ spl24_124 | spl24_259)), inference(forward_demodulation, [], [f14090, f1152])).
fof(f14090, plain, (~ (e1 = op(op(e0, e3), e0)) | (~ spl24_124 | spl24_259)), inference(forward_demodulation, [], [f1940, f1211])).
fof(f14080, plain, (~ spl24_118 | ~ spl24_48 | ~ spl24_194), inference(avatar_split_clause, [], [f14079, f1597, f890, f1184])).
fof(f14079, plain, (~ (e2 = op(e0, e1)) | (~ spl24_48 | ~ spl24_194)), inference(forward_demodulation, [], [f13448, f892])).
fof(f13448, plain, (~ (op(e0, e1) = op(e3, e0)) | ~ spl24_194), inference(backward_demodulation, [], [f387, f1598])).
fof(f387, plain, ~ (op(e1, e0) = op(e3, e0)), inference(cnf_transformation, [], [f4])).
fof(f14078, plain, (~ spl24_119 | ~ spl24_124), inference(avatar_split_clause, [], [f14077, f1209, f1188])).
fof(f14077, plain, (~ (e3 = op(e0, e1)) | ~ spl24_124), inference(forward_demodulation, [], [f432, f1211])).
fof(f432, plain, ~ (op(e0, e0) = op(e0, e1)), inference(cnf_transformation, [], [f4])).
fof(f14069, plain, (~ spl24_98 | ~ spl24_112 | spl24_202), inference(avatar_split_clause, [], [f13974, f1636, f1159, f1100])).
fof(f1636, plain, (spl24_202 <=> (e2 = op(op(e0, e2), e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_202])])).
fof(f13974, plain, (~ (e2 = op(e1, e0)) | (~ spl24_112 | spl24_202)), inference(backward_demodulation, [], [f1638, f1161])).
fof(f1638, plain, (~ (e2 = op(op(e0, e2), e0)) | spl24_202), inference(avatar_component_clause, [], [f1636])).
fof(f14062, plain, (~ spl24_83 | ~ spl24_48 | ~ spl24_159), inference(avatar_split_clause, [], [f14061, f1424, f890, f1037])).
fof(f14061, plain, (~ (e2 = op(e1, e3)) | (~ spl24_48 | ~ spl24_159)), inference(forward_demodulation, [], [f13468, f892])).
fof(f13468, plain, (~ (op(e1, e3) = op(e3, e0)) | ~ spl24_159), inference(backward_demodulation, [], [f462, f1425])).
fof(f14055, plain, (~ spl24_70 | ~ spl24_56 | ~ spl24_64 | ~ spl24_112 | spl24_266), inference(avatar_split_clause, [], [f13978, f2000, f1159, f957, f924, f982])).
fof(f2000, plain, (spl24_266 <=> (e4 = op(e2, op(op(e2, op(e2, e2)), e2)))), introduced(avatar_definition, [new_symbols(naming, [spl24_266])])).
fof(f13978, plain, (~ (e4 = op(e2, e1)) | (~ spl24_56 | ~ spl24_64 | ~ spl24_112 | spl24_266)), inference(backward_demodulation, [], [f13939, f1161])).
fof(f13939, plain, (~ (e4 = op(e2, op(e0, e2))) | (~ spl24_56 | ~ spl24_64 | spl24_266)), inference(forward_demodulation, [], [f13938, f926])).
fof(f13938, plain, (~ (e4 = op(e2, op(op(e2, e3), e2))) | (~ spl24_64 | spl24_266)), inference(forward_demodulation, [], [f2002, f959])).
fof(f2002, plain, (~ (e4 = op(e2, op(op(e2, op(e2, e2)), e2))) | spl24_266), inference(avatar_component_clause, [], [f2000])).
fof(f14051, plain, (~ spl24_45 | ~ spl24_32 | spl24_212), inference(avatar_split_clause, [], [f13996, f1688, f823, f877])).
fof(f13996, plain, (~ (e4 = op(e3, e1)) | (~ spl24_32 | spl24_212)), inference(backward_demodulation, [], [f1690, f825])).
fof(f14049, plain, (~ spl24_41 | ~ spl24_32 | spl24_249), inference(avatar_split_clause, [], [f14001, f1871, f823, f861])).
fof(f1871, plain, (spl24_249 <=> (e0 = op(e3, op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl24_249])])).
fof(f14001, plain, (~ (e0 = op(e3, e1)) | (~ spl24_32 | spl24_249)), inference(backward_demodulation, [], [f1873, f825])).
fof(f1873, plain, (~ (e0 = op(e3, op(e3, e3))) | spl24_249), inference(avatar_component_clause, [], [f1871])).
fof(f14042, plain, (~ spl24_14 | ~ spl24_64), inference(avatar_split_clause, [], [f14041, f957, f747])).
fof(f14041, plain, (~ (e3 = op(e4, e2)) | ~ spl24_64), inference(forward_demodulation, [], [f410, f959])).
fof(f14021, plain, (~ spl24_11 | ~ spl24_103 | spl24_136), inference(avatar_contradiction_clause, [], [f14020])).
fof(f14020, plain, ($false | (~ spl24_11 | ~ spl24_103 | spl24_136)), inference(subsumption_resolution, [], [f14019, f1123])).
fof(f14019, plain, (~ (e2 = op(e0, e4)) | (~ spl24_11 | spl24_136)), inference(forward_demodulation, [], [f1314, f737])).
fof(f14018, plain, (~ spl24_11 | ~ spl24_112 | ~ spl24_137), inference(avatar_contradiction_clause, [], [f14017])).
fof(f14017, plain, ($false | (~ spl24_11 | ~ spl24_112 | ~ spl24_137)), inference(subsumption_resolution, [], [f14016, f488])).
fof(f14016, plain, ((e1 = e4) | (~ spl24_11 | ~ spl24_112 | ~ spl24_137)), inference(forward_demodulation, [], [f14015, f1161])).
fof(f14015, plain, ((e4 = op(e0, e2)) | (~ spl24_11 | ~ spl24_137)), inference(forward_demodulation, [], [f1319, f737])).
fof(f13991, plain, (~ spl24_35 | ~ spl24_40), inference(avatar_split_clause, [], [f13988, f856, f835])).
fof(f13988, plain, (~ (e4 = op(e3, e3)) | ~ spl24_40), inference(backward_demodulation, [], [f469, f858])).
fof(f469, plain, ~ (op(e3, e2) = op(e3, e3)), inference(cnf_transformation, [], [f4])).
fof(f13980, plain, (~ spl24_107 | ~ spl24_112), inference(avatar_split_clause, [], [f13970, f1159, f1138])).
fof(f13970, plain, (~ (e1 = op(e0, e3)) | ~ spl24_112), inference(backward_demodulation, [], [f439, f1161])).
fof(f13979, plain, (~ spl24_37 | ~ spl24_112), inference(avatar_split_clause, [], [f13969, f1159, f844])).
fof(f13969, plain, (~ (e1 = op(e3, e2)) | ~ spl24_112), inference(backward_demodulation, [], [f404, f1161])).
fof(f13913, plain, (~ spl24_115 | ~ spl24_56 | ~ spl24_64 | spl24_221), inference(avatar_split_clause, [], [f13882, f1728, f957, f924, f1171])).
fof(f1728, plain, (spl24_221 <=> (e4 = op(op(e2, op(e2, e2)), e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_221])])).
fof(f13882, plain, (~ (e4 = op(e0, e2)) | (~ spl24_56 | ~ spl24_64 | spl24_221)), inference(forward_demodulation, [], [f13872, f926])).
fof(f13872, plain, (~ (e4 = op(op(e2, e3), e2)) | (~ spl24_64 | spl24_221)), inference(backward_demodulation, [], [f1730, f959])).
fof(f1730, plain, (~ (e4 = op(op(e2, op(e2, e2)), e2)) | spl24_221), inference(avatar_component_clause, [], [f1728])).
fof(f13908, plain, (~ spl24_108 | ~ spl24_48 | spl24_163), inference(avatar_split_clause, [], [f13907, f1443, f890, f1142])).
fof(f13907, plain, (~ (e2 = op(e0, e3)) | (~ spl24_48 | spl24_163)), inference(forward_demodulation, [], [f1445, f892])).
fof(f13904, plain, (spl24_108 | ~ spl24_56 | ~ spl24_169), inference(avatar_split_clause, [], [f13842, f1472, f924, f1142])).
fof(f13842, plain, ((e2 = op(e0, e3)) | (~ spl24_56 | ~ spl24_169)), inference(forward_demodulation, [], [f1474, f926])).
fof(f13891, plain, (~ spl24_52 | ~ spl24_56 | ~ spl24_64 | ~ spl24_115 | spl24_240), inference(avatar_contradiction_clause, [], [f13890])).
fof(f13890, plain, ($false | (~ spl24_52 | ~ spl24_56 | ~ spl24_64 | ~ spl24_115 | spl24_240)), inference(subsumption_resolution, [], [f13889, f909])).
fof(f13889, plain, (~ (e1 = op(e2, e4)) | (~ spl24_56 | ~ spl24_64 | ~ spl24_115 | spl24_240)), inference(forward_demodulation, [], [f13888, f1173])).
fof(f13888, plain, (~ (e1 = op(e2, op(e0, e2))) | (~ spl24_56 | ~ spl24_64 | spl24_240)), inference(forward_demodulation, [], [f13876, f926])).
fof(f13876, plain, (~ (e1 = op(e2, op(op(e2, e3), e2))) | (~ spl24_64 | spl24_240)), inference(backward_demodulation, [], [f1828, f959])).
fof(f1828, plain, (~ (e1 = op(e2, op(op(e2, op(e2, e2)), e2))) | spl24_240), inference(avatar_component_clause, [], [f1826])).
fof(f1826, plain, (spl24_240 <=> (e1 = op(e2, op(op(e2, op(e2, e2)), e2)))), introduced(avatar_definition, [new_symbols(naming, [spl24_240])])).
fof(f13858, plain, (~ spl24_74 | ~ spl24_124), inference(avatar_split_clause, [], [f13854, f1209, f999])).
fof(f999, plain, (spl24_74 <=> (e3 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_74])])).
fof(f13854, plain, (~ (e3 = op(e2, e0)) | ~ spl24_124), inference(backward_demodulation, [], [f383, f1211])).
fof(f13849, plain, (~ spl24_74 | ~ spl24_48 | spl24_162), inference(avatar_split_clause, [], [f13848, f1438, f890, f999])).
fof(f1438, plain, (spl24_162 <=> (e3 = op(op(e3, e0), e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_162])])).
fof(f13848, plain, (~ (e3 = op(e2, e0)) | (~ spl24_48 | spl24_162)), inference(forward_demodulation, [], [f1439, f892])).
fof(f1439, plain, (~ (e3 = op(op(e3, e0), e0)) | spl24_162), inference(avatar_component_clause, [], [f1438])).
fof(f13847, plain, (~ spl24_114 | ~ spl24_56 | spl24_168), inference(avatar_split_clause, [], [f13846, f1467, f924, f1167])).
fof(f13846, plain, (~ (e3 = op(e0, e2)) | (~ spl24_56 | spl24_168)), inference(forward_demodulation, [], [f1469, f926])).
fof(f13712, plain, (spl24_26 | ~ spl24_35 | ~ spl24_249), inference(avatar_split_clause, [], [f13711, f1871, f835, f798])).
fof(f13711, plain, ((e0 = op(e3, e4)) | (~ spl24_35 | ~ spl24_249)), inference(backward_demodulation, [], [f1872, f837])).
fof(f1872, plain, ((e0 = op(e3, op(e3, e3))) | ~ spl24_249), inference(avatar_component_clause, [], [f1871])).
fof(f13698, plain, (~ spl24_48 | ~ spl24_56 | spl24_161), inference(avatar_contradiction_clause, [], [f13697])).
fof(f13697, plain, ($false | (~ spl24_48 | ~ spl24_56 | spl24_161)), inference(subsumption_resolution, [], [f13695, f926])).
fof(f13695, plain, (~ (e0 = op(e2, e3)) | (~ spl24_48 | spl24_161)), inference(backward_demodulation, [], [f1435, f892])).
fof(f13657, plain, (~ spl24_20 | ~ spl24_80 | spl24_181), inference(avatar_split_clause, [], [f13650, f1531, f1024, f772])).
fof(f13650, plain, (~ (e4 = op(e4, e1)) | (~ spl24_80 | spl24_181)), inference(backward_demodulation, [], [f1533, f1026])).
fof(f13604, plain, (~ spl24_77 | ~ spl24_92), inference(avatar_split_clause, [], [f13597, f1075, f1012])).
fof(f13597, plain, (~ (e1 = op(e1, e4)) | ~ spl24_92), inference(backward_demodulation, [], [f448, f1077])).
fof(f13587, plain, (~ spl24_75 | ~ spl24_103 | spl24_196), inference(avatar_split_clause, [], [f13584, f1606, f1121, f1003])).
fof(f13584, plain, (~ (e4 = op(e2, e0)) | (~ spl24_103 | spl24_196)), inference(backward_demodulation, [], [f1608, f1123])).
fof(f13586, plain, (~ spl24_53 | ~ spl24_103), inference(avatar_split_clause, [], [f13578, f1121, f911])).
fof(f13578, plain, (~ (e2 = op(e2, e4)) | ~ spl24_103), inference(backward_demodulation, [], [f423, f1123])).
fof(f13574, plain, (~ spl24_57 | ~ spl24_107), inference(avatar_split_clause, [], [f13561, f1138, f928])).
fof(f13561, plain, (~ (e1 = op(e2, e3)) | ~ spl24_107), inference(backward_demodulation, [], [f413, f1140])).
fof(f13558, plain, (~ spl24_107 | ~ spl24_114 | ~ spl24_179 | spl24_243 | ~ spl24_252), inference(avatar_split_clause, [], [f13551, f1888, f1844, f1522, f1167, f1138])).
fof(f13551, plain, (~ (e1 = op(e0, e3)) | (~ spl24_114 | ~ spl24_179 | spl24_243 | ~ spl24_252)), inference(backward_demodulation, [], [f13462, f1169])).
fof(f13462, plain, (~ (e1 = op(e0, op(e0, e2))) | (~ spl24_179 | spl24_243 | ~ spl24_252)), inference(backward_demodulation, [], [f13238, f1523])).
fof(f13238, plain, (~ (e1 = op(e0, op(e2, e0))) | (spl24_243 | ~ spl24_252)), inference(backward_demodulation, [], [f1846, f1889])).
fof(f13536, plain, (~ spl24_21 | ~ spl24_116 | ~ spl24_194), inference(avatar_split_clause, [], [f13529, f1597, f1176, f777])).
fof(f13529, plain, (~ (e0 = op(e4, e0)) | (~ spl24_116 | ~ spl24_194)), inference(backward_demodulation, [], [f13449, f1178])).
fof(f13449, plain, (~ (op(e0, e1) = op(e4, e0)) | ~ spl24_194), inference(backward_demodulation, [], [f388, f1598])).
fof(f13534, plain, (~ spl24_101 | ~ spl24_116), inference(avatar_split_clause, [], [f13519, f1176, f1113])).
fof(f13519, plain, (~ (e0 = op(e0, e4)) | ~ spl24_116), inference(backward_demodulation, [], [f438, f1178])).
fof(f13533, plain, (~ spl24_91 | ~ spl24_116), inference(avatar_split_clause, [], [f13513, f1176, f1071])).
fof(f13513, plain, (~ (e0 = op(e1, e1)) | ~ spl24_116), inference(backward_demodulation, [], [f392, f1178])).
fof(f13512, plain, (spl24_103 | ~ spl24_125 | ~ spl24_252), inference(avatar_split_clause, [], [f13508, f1888, f1213, f1121])).
fof(f13508, plain, ((e2 = op(e0, e4)) | (~ spl24_125 | ~ spl24_252)), inference(backward_demodulation, [], [f1889, f1215])).
fof(f13511, plain, (~ spl24_120 | ~ spl24_125), inference(avatar_split_clause, [], [f13504, f1213, f1192])).
fof(f13504, plain, (~ (e4 = op(e0, e1)) | ~ spl24_125), inference(backward_demodulation, [], [f432, f1215])).
fof(f13442, plain, (~ spl24_69 | spl24_228 | ~ spl24_236), inference(avatar_split_clause, [], [f13441, f1800, f1760, f978])).
fof(f13441, plain, (~ (e3 = op(e2, e1)) | (spl24_228 | ~ spl24_236)), inference(forward_demodulation, [], [f1762, f1801])).
fof(f13429, plain, (spl24_74 | ~ spl24_244 | ~ spl24_252), inference(avatar_split_clause, [], [f13428, f1888, f1848, f999])).
fof(f13428, plain, ((e3 = op(e2, e0)) | (~ spl24_244 | ~ spl24_252)), inference(forward_demodulation, [], [f1849, f1889])).
fof(f13427, plain, (~ spl24_75 | spl24_247 | ~ spl24_252), inference(avatar_split_clause, [], [f13426, f1888, f1862, f1003])).
fof(f1862, plain, (spl24_247 <=> (e4 = op(op(e0, op(e0, e0)), e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_247])])).
fof(f13426, plain, (~ (e4 = op(e2, e0)) | (spl24_247 | ~ spl24_252)), inference(forward_demodulation, [], [f1864, f1889])).
fof(f1864, plain, (~ (e4 = op(op(e0, op(e0, e0)), e0)) | spl24_247), inference(avatar_component_clause, [], [f1862])).
fof(f13385, plain, (~ spl24_69 | ~ spl24_85 | ~ spl24_91 | ~ spl24_98 | spl24_268), inference(avatar_contradiction_clause, [], [f13384])).
fof(f13384, plain, ($false | (~ spl24_69 | ~ spl24_85 | ~ spl24_91 | ~ spl24_98 | spl24_268)), inference(subsumption_resolution, [], [f13383, f1047])).
fof(f13383, plain, (~ (e4 = op(e1, e3)) | (~ spl24_69 | ~ spl24_91 | ~ spl24_98 | spl24_268)), inference(forward_demodulation, [], [f13382, f980])).
fof(f13382, plain, (~ (e4 = op(e1, op(e2, e1))) | (~ spl24_91 | ~ spl24_98 | spl24_268)), inference(forward_demodulation, [], [f13381, f1102])).
fof(f13381, plain, (~ (e4 = op(e1, op(op(e1, e0), e1))) | (~ spl24_91 | spl24_268)), inference(forward_demodulation, [], [f2012, f1073])).
fof(f2012, plain, (~ (e4 = op(e1, op(op(e1, op(e1, e1)), e1))) | spl24_268), inference(avatar_component_clause, [], [f2010])).
fof(f2010, plain, (spl24_268 <=> (e4 = op(e1, op(op(e1, op(e1, e1)), e1)))), introduced(avatar_definition, [new_symbols(naming, [spl24_268])])).
fof(f13374, plain, (~ spl24_37 | ~ spl24_47), inference(avatar_split_clause, [], [f13365, f886, f844])).
fof(f13365, plain, (~ (e1 = op(e3, e2)) | ~ spl24_47), inference(backward_demodulation, [], [f463, f888])).
fof(f13344, plain, (~ spl24_60 | ~ spl24_75), inference(avatar_split_clause, [], [f13331, f1003, f940])).
fof(f13331, plain, (~ (e4 = op(e2, e3)) | ~ spl24_75), inference(backward_demodulation, [], [f454, f1005])).
fof(f13329, plain, (spl24_7 | ~ spl24_85 | ~ spl24_185), inference(avatar_contradiction_clause, [], [f13328])).
fof(f13328, plain, ($false | (spl24_7 | ~ spl24_85 | ~ spl24_185)), inference(subsumption_resolution, [], [f13325, f719])).
fof(f719, plain, (~ (e1 = op(e4, e3)) | spl24_7), inference(avatar_component_clause, [], [f718])).
fof(f13325, plain, ((e1 = op(e4, e3)) | (~ spl24_85 | ~ spl24_185)), inference(backward_demodulation, [], [f1553, f1047])).
fof(f1553, plain, ((e1 = op(op(e1, e3), e3)) | ~ spl24_185), inference(avatar_component_clause, [], [f1551])).
fof(f1551, plain, (spl24_185 <=> (e1 = op(op(e1, e3), e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_185])])).
fof(f13296, plain, (~ spl24_72 | ~ spl24_112 | spl24_179), inference(avatar_split_clause, [], [f13291, f1522, f1159, f991])).
fof(f13291, plain, (~ (e1 = op(e2, e0)) | (~ spl24_112 | spl24_179)), inference(backward_demodulation, [], [f1524, f1161])).
fof(f13260, plain, (~ spl24_83 | ~ spl24_91 | ~ spl24_228 | spl24_256), inference(avatar_split_clause, [], [f13259, f1914, f1760, f1071, f1037])).
fof(f1914, plain, (spl24_256 <=> (e2 = op(e1, op(op(e1, op(e1, e1)), e1)))), introduced(avatar_definition, [new_symbols(naming, [spl24_256])])).
fof(f13259, plain, (~ (e2 = op(e1, e3)) | (~ spl24_91 | ~ spl24_228 | spl24_256)), inference(backward_demodulation, [], [f13136, f13256])).
fof(f13256, plain, ((e3 = op(op(e1, e0), e1)) | (~ spl24_91 | ~ spl24_228)), inference(forward_demodulation, [], [f1761, f1073])).
fof(f1761, plain, ((e3 = op(op(e1, op(e1, e1)), e1)) | ~ spl24_228), inference(avatar_component_clause, [], [f1760])).
fof(f13136, plain, (~ (e2 = op(e1, op(op(e1, e0), e1))) | (~ spl24_91 | spl24_256)), inference(forward_demodulation, [], [f1916, f1073])).
fof(f1916, plain, (~ (e2 = op(e1, op(op(e1, op(e1, e1)), e1))) | spl24_256), inference(avatar_component_clause, [], [f1914])).
fof(f13227, plain, (~ spl24_110 | ~ spl24_120), inference(avatar_split_clause, [], [f13226, f1192, f1150])).
fof(f13226, plain, (~ (e4 = op(e0, e3)) | ~ spl24_120), inference(forward_demodulation, [], [f437, f1194])).
fof(f13223, plain, (~ spl24_100 | ~ spl24_120 | spl24_194), inference(avatar_split_clause, [], [f13222, f1597, f1192, f1108])).
fof(f13222, plain, (~ (e4 = op(e1, e0)) | (~ spl24_120 | spl24_194)), inference(forward_demodulation, [], [f1599, f1194])).
fof(f13215, plain, (~ spl24_83 | ~ spl24_43 | spl24_159), inference(avatar_split_clause, [], [f13214, f1424, f869, f1037])).
fof(f13214, plain, (~ (e2 = op(e1, e3)) | (~ spl24_43 | spl24_159)), inference(forward_demodulation, [], [f1426, f871])).
fof(f13207, plain, (~ spl24_47 | ~ spl24_31 | spl24_233), inference(avatar_split_clause, [], [f13045, f1783, f819, f886])).
fof(f13045, plain, (~ (e1 = op(e3, e0)) | (~ spl24_31 | spl24_233)), inference(backward_demodulation, [], [f1785, f821])).
fof(f1785, plain, (~ (e1 = op(e3, op(e3, e3))) | spl24_233), inference(avatar_component_clause, [], [f1783])).
fof(f13204, plain, (~ spl24_38 | ~ spl24_43), inference(avatar_split_clause, [], [f13203, f869, f848])).
fof(f13203, plain, (~ (e2 = op(e3, e2)) | ~ spl24_43), inference(forward_demodulation, [], [f466, f871])).
fof(f466, plain, ~ (op(e3, e1) = op(e3, e2)), inference(cnf_transformation, [], [f4])).
fof(f13202, plain, (~ spl24_43 | ~ spl24_90 | ~ spl24_232 | spl24_268), inference(avatar_contradiction_clause, [], [f13201])).
fof(f13201, plain, ($false | (~ spl24_43 | ~ spl24_90 | ~ spl24_232 | spl24_268)), inference(subsumption_resolution, [], [f13198, f1068])).
fof(f13198, plain, (~ (e4 = op(e1, e2)) | (~ spl24_43 | ~ spl24_232 | spl24_268)), inference(backward_demodulation, [], [f12811, f871])).
fof(f12811, plain, (~ (e4 = op(e1, op(e3, e1))) | (~ spl24_232 | spl24_268)), inference(backward_demodulation, [], [f2012, f1779])).
fof(f13168, plain, (~ spl24_118 | ~ spl24_123), inference(avatar_split_clause, [], [f13161, f1205, f1184])).
fof(f13161, plain, (~ (e2 = op(e0, e1)) | ~ spl24_123), inference(backward_demodulation, [], [f432, f1207])).
fof(f13167, plain, (~ spl24_48 | ~ spl24_123), inference(avatar_split_clause, [], [f13160, f1205, f890])).
fof(f13160, plain, (~ (e2 = op(e3, e0)) | ~ spl24_123), inference(backward_demodulation, [], [f384, f1207])).
fof(f13146, plain, (~ spl24_48 | ~ spl24_248 | spl24_251), inference(avatar_split_clause, [], [f13141, f1882, f1866, f890])).
fof(f13141, plain, (~ (e2 = op(e3, e0)) | (~ spl24_248 | spl24_251)), inference(backward_demodulation, [], [f1884, f1867])).
fof(f13105, plain, (~ spl24_119 | ~ spl24_114), inference(avatar_split_clause, [], [f13104, f1167, f1188])).
fof(f13104, plain, (~ (e3 = op(e0, e1)) | ~ spl24_114), inference(forward_demodulation, [], [f436, f1169])).
fof(f13103, plain, (~ spl24_119 | ~ spl24_99 | spl24_194), inference(avatar_split_clause, [], [f13102, f1597, f1104, f1188])).
fof(f13102, plain, (~ (e3 = op(e0, e1)) | (~ spl24_99 | spl24_194)), inference(forward_demodulation, [], [f1599, f1106])).
fof(f13098, plain, (~ spl24_50 | ~ spl24_31 | spl24_212), inference(avatar_split_clause, [], [f13041, f1688, f819, f898])).
fof(f13041, plain, (~ (e4 = op(e3, e0)) | (~ spl24_31 | spl24_212)), inference(backward_demodulation, [], [f1690, f821])).
fof(f13093, plain, (~ spl24_43 | ~ spl24_232 | spl24_235), inference(avatar_split_clause, [], [f12809, f1794, f1778, f869])).
fof(f1794, plain, (spl24_235 <=> (e2 = op(op(e1, op(e1, e1)), e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_235])])).
fof(f12809, plain, (~ (e2 = op(e3, e1)) | (~ spl24_232 | spl24_235)), inference(backward_demodulation, [], [f1796, f1779])).
fof(f1796, plain, (~ (e2 = op(op(e1, op(e1, e1)), e1)) | spl24_235), inference(avatar_component_clause, [], [f1794])).
fof(f13056, plain, (~ spl24_31 | ~ spl24_48 | spl24_224), inference(avatar_contradiction_clause, [], [f13055])).
fof(f13055, plain, ($false | (~ spl24_31 | ~ spl24_48 | spl24_224)), inference(subsumption_resolution, [], [f13043, f892])).
fof(f13043, plain, (~ (e2 = op(e3, e0)) | (~ spl24_31 | spl24_224)), inference(backward_demodulation, [], [f1743, f821])).
fof(f1743, plain, (~ (e2 = op(e3, op(e3, e3))) | spl24_224), inference(avatar_component_clause, [], [f1741])).
fof(f12993, plain, (~ spl24_14 | ~ spl24_60 | spl24_168), inference(avatar_split_clause, [], [f12990, f1467, f940, f747])).
fof(f12990, plain, (~ (e3 = op(e4, e2)) | (~ spl24_60 | spl24_168)), inference(backward_demodulation, [], [f1469, f942])).
fof(f12992, plain, (~ spl24_10 | ~ spl24_60), inference(avatar_split_clause, [], [f12987, f940, f730])).
fof(f12987, plain, (~ (e4 = op(e4, e3)) | ~ spl24_60), inference(backward_demodulation, [], [f420, f942])).
fof(f12923, plain, (~ spl24_74 | ~ spl24_99), inference(avatar_split_clause, [], [f12915, f1104, f999])).
fof(f12915, plain, (~ (e3 = op(e2, e0)) | ~ spl24_99), inference(backward_demodulation, [], [f386, f1106])).
fof(f12912, plain, (~ spl24_76 | ~ spl24_101), inference(avatar_split_clause, [], [f12895, f1113, f1008])).
fof(f12895, plain, (~ (e0 = op(e1, e4)) | ~ spl24_101), inference(backward_demodulation, [], [f422, f1115])).
fof(f1115, plain, ((e0 = op(e0, e4)) | ~ spl24_101), inference(avatar_component_clause, [], [f1113])).
fof(f12881, plain, (spl24_36 | ~ spl24_114 | ~ spl24_203), inference(avatar_contradiction_clause, [], [f12880])).
fof(f12880, plain, ($false | (spl24_36 | ~ spl24_114 | ~ spl24_203)), inference(subsumption_resolution, [], [f12875, f841])).
fof(f841, plain, (~ (e0 = op(e3, e2)) | spl24_36), inference(avatar_component_clause, [], [f840])).
fof(f840, plain, (spl24_36 <=> (e0 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_36])])).
fof(f12875, plain, ((e0 = op(e3, e2)) | (~ spl24_114 | ~ spl24_203)), inference(backward_demodulation, [], [f1643, f1169])).
fof(f12879, plain, (~ spl24_48 | ~ spl24_114 | spl24_202), inference(avatar_split_clause, [], [f12874, f1636, f1167, f890])).
fof(f12874, plain, (~ (e2 = op(e3, e0)) | (~ spl24_114 | spl24_202)), inference(backward_demodulation, [], [f1638, f1169])).
fof(f12878, plain, (~ spl24_74 | ~ spl24_114 | spl24_179), inference(avatar_split_clause, [], [f12873, f1522, f1167, f999])).
fof(f12873, plain, (~ (e3 = op(e2, e0)) | (~ spl24_114 | spl24_179)), inference(backward_demodulation, [], [f1524, f1169])).
fof(f12877, plain, (~ spl24_104 | ~ spl24_114), inference(avatar_split_clause, [], [f12871, f1167, f1125])).
fof(f12871, plain, (~ (e3 = op(e0, e4)) | ~ spl24_114), inference(backward_demodulation, [], [f440, f1169])).
fof(f12876, plain, (~ spl24_14 | ~ spl24_114), inference(avatar_split_clause, [], [f12869, f1167, f747])).
fof(f12869, plain, (~ (e3 = op(e4, e2)) | ~ spl24_114), inference(backward_demodulation, [], [f405, f1169])).
fof(f12831, plain, (~ spl24_70 | ~ spl24_90 | spl24_175), inference(avatar_split_clause, [], [f12830, f1503, f1066, f982])).
fof(f12830, plain, (~ (e4 = op(e2, e1)) | (~ spl24_90 | spl24_175)), inference(forward_demodulation, [], [f1505, f1068])).
fof(f12812, plain, (~ spl24_45 | spl24_231 | ~ spl24_232), inference(avatar_split_clause, [], [f12808, f1778, f1774, f877])).
fof(f12808, plain, (~ (e4 = op(e3, e1)) | (spl24_231 | ~ spl24_232)), inference(backward_demodulation, [], [f1776, f1779])).
fof(f12804, plain, (~ spl24_104 | ~ spl24_125 | spl24_248), inference(avatar_split_clause, [], [f12803, f1866, f1213, f1125])).
fof(f12803, plain, (~ (e3 = op(e0, e4)) | (~ spl24_125 | spl24_248)), inference(forward_demodulation, [], [f1868, f1215])).
fof(f12743, plain, (~ spl24_69 | ~ spl24_74), inference(avatar_split_clause, [], [f12737, f999, f978])).
fof(f12737, plain, (~ (e3 = op(e2, e1)) | ~ spl24_74), inference(backward_demodulation, [], [f452, f1001])).
fof(f1001, plain, ((e3 = op(e2, e0)) | ~ spl24_74), inference(avatar_component_clause, [], [f999])).
fof(f12710, plain, (~ spl24_79 | ~ spl24_94), inference(avatar_split_clause, [], [f12699, f1083, f1020])).
fof(f12699, plain, (~ (e3 = op(e1, e4)) | ~ spl24_94), inference(backward_demodulation, [], [f448, f1085])).
fof(f12709, plain, (~ spl24_69 | ~ spl24_94), inference(avatar_split_clause, [], [f12697, f1083, f978])).
fof(f12697, plain, (~ (e3 = op(e2, e1)) | ~ spl24_94), inference(backward_demodulation, [], [f396, f1085])).
fof(f12686, plain, (~ spl24_18 | ~ spl24_118), inference(avatar_split_clause, [], [f12680, f1184, f764])).
fof(f12680, plain, (~ (e2 = op(e4, e1)) | ~ spl24_118), inference(backward_demodulation, [], [f395, f1186])).
fof(f12677, plain, (~ spl24_103 | ~ spl24_125 | spl24_252), inference(avatar_split_clause, [], [f12672, f1888, f1213, f1121])).
fof(f12672, plain, (~ (e2 = op(e0, e4)) | (~ spl24_125 | spl24_252)), inference(backward_demodulation, [], [f1890, f1215])).
fof(f12659, plain, (spl24_70 | ~ spl24_90 | ~ spl24_175), inference(avatar_split_clause, [], [f12658, f1503, f1066, f982])).
fof(f12658, plain, ((e4 = op(e2, e1)) | (~ spl24_90 | ~ spl24_175)), inference(forward_demodulation, [], [f1504, f1068])).
fof(f1504, plain, ((op(e1, e2) = op(e2, e1)) | ~ spl24_175), inference(avatar_component_clause, [], [f1503])).
fof(f12657, plain, (~ spl24_18 | ~ spl24_90 | spl24_187), inference(avatar_split_clause, [], [f12656, f1561, f1066, f764])).
fof(f12656, plain, (~ (e2 = op(e4, e1)) | (~ spl24_90 | spl24_187)), inference(forward_demodulation, [], [f1563, f1068])).
fof(f12653, plain, (~ spl24_124 | ~ spl24_106 | spl24_199), inference(avatar_split_clause, [], [f12652, f1621, f1134, f1209])).
fof(f12652, plain, (~ (op(e0, e0) = e3) | (~ spl24_106 | spl24_199)), inference(forward_demodulation, [], [f1623, f1136])).
fof(f12628, plain, (~ spl24_93 | ~ spl24_98), inference(avatar_split_clause, [], [f12627, f1100, f1079])).
fof(f12627, plain, (~ (e2 = op(e1, e1)) | ~ spl24_98), inference(forward_demodulation, [], [f442, f1102])).
fof(f12625, plain, (~ spl24_76 | ~ spl24_22 | spl24_144), inference(avatar_split_clause, [], [f12468, f1350, f781, f1008])).
fof(f12468, plain, (~ (e0 = op(e1, e4)) | (~ spl24_22 | spl24_144)), inference(backward_demodulation, [], [f1352, f783])).
fof(f12600, plain, (~ spl24_14 | ~ spl24_38 | ~ spl24_137), inference(avatar_contradiction_clause, [], [f12599])).
fof(f12599, plain, ($false | (~ spl24_14 | ~ spl24_38 | ~ spl24_137)), inference(subsumption_resolution, [], [f12598, f490])).
fof(f12598, plain, ((e2 = e4) | (~ spl24_14 | ~ spl24_38 | ~ spl24_137)), inference(forward_demodulation, [], [f12597, f850])).
fof(f850, plain, ((e2 = op(e3, e2)) | ~ spl24_38), inference(avatar_component_clause, [], [f848])).
fof(f12597, plain, ((e4 = op(e3, e2)) | (~ spl24_14 | ~ spl24_137)), inference(forward_demodulation, [], [f1319, f749])).
fof(f12538, plain, (~ spl24_22 | spl24_100 | ~ spl24_145), inference(avatar_contradiction_clause, [], [f12537])).
fof(f12537, plain, ($false | (~ spl24_22 | spl24_100 | ~ spl24_145)), inference(subsumption_resolution, [], [f12536, f1109])).
fof(f1109, plain, (~ (e4 = op(e1, e0)) | spl24_100), inference(avatar_component_clause, [], [f1108])).
fof(f12536, plain, ((e4 = op(e1, e0)) | (~ spl24_22 | ~ spl24_145)), inference(forward_demodulation, [], [f1357, f783])).
fof(f12490, plain, (~ spl24_80 | ~ spl24_90), inference(avatar_split_clause, [], [f12489, f1066, f1024])).
fof(f12489, plain, (~ (e4 = op(e1, e4)) | ~ spl24_90), inference(forward_demodulation, [], [f450, f1068])).
fof(f12474, plain, (~ spl24_12 | ~ spl24_22), inference(avatar_split_clause, [], [f12464, f781, f739])).
fof(f12464, plain, (~ (e1 = op(e4, e2)) | ~ spl24_22), inference(backward_demodulation, [], [f473, f783])).
fof(f473, plain, ~ (op(e4, e0) = op(e4, e2)), inference(cnf_transformation, [], [f4])).
fof(f12399, plain, (~ spl24_102 | ~ spl24_112), inference(avatar_split_clause, [], [f12389, f1159, f1117])).
fof(f12389, plain, (~ (e1 = op(e0, e4)) | ~ spl24_112), inference(backward_demodulation, [], [f440, f1161])).
fof(f12386, plain, (~ spl24_22 | ~ spl24_120 | spl24_205), inference(avatar_split_clause, [], [f12381, f1651, f1192, f781])).
fof(f12381, plain, (~ (e1 = op(e4, e0)) | (~ spl24_120 | spl24_205)), inference(backward_demodulation, [], [f1653, f1194])).
fof(f12361, plain, (spl24_99 | ~ spl24_91 | ~ spl24_232), inference(avatar_split_clause, [], [f12360, f1778, f1071, f1104])).
fof(f12360, plain, ((e3 = op(e1, e0)) | (~ spl24_91 | ~ spl24_232)), inference(forward_demodulation, [], [f1779, f1073])).
fof(f12354, plain, (~ spl24_115 | ~ spl24_123 | spl24_245), inference(avatar_split_clause, [], [f12353, f1852, f1205, f1171])).
fof(f12353, plain, (~ (e4 = op(e0, e2)) | (~ spl24_123 | spl24_245)), inference(forward_demodulation, [], [f1854, f1207])).
fof(f1854, plain, (~ (e4 = op(e0, op(e0, e0))) | spl24_245), inference(avatar_component_clause, [], [f1852])).
fof(f12350, plain, (spl24_112 | ~ spl24_123 | ~ spl24_260), inference(avatar_split_clause, [], [f12349, f1944, f1205, f1159])).
fof(f12349, plain, ((e1 = op(e0, e2)) | (~ spl24_123 | ~ spl24_260)), inference(forward_demodulation, [], [f1945, f1207])).
fof(f12333, plain, (~ spl24_75 | ~ spl24_61 | spl24_218), inference(avatar_split_clause, [], [f12332, f1714, f945, f1003])).
fof(f12332, plain, (~ (e4 = op(e2, e0)) | (~ spl24_61 | spl24_218)), inference(forward_demodulation, [], [f1716, f947])).
fof(f12292, plain, (~ spl24_52 | ~ spl24_18 | spl24_140), inference(avatar_split_clause, [], [f11795, f1331, f764, f907])).
fof(f11795, plain, (~ (e1 = op(e2, e4)) | (~ spl24_18 | spl24_140)), inference(forward_demodulation, [], [f1333, f766])).
fof(f12288, plain, (~ spl24_24 | ~ spl24_1 | spl24_215), inference(avatar_split_clause, [], [f12195, f1701, f693, f789])).
fof(f12195, plain, (~ (e3 = op(e4, e0)) | (~ spl24_1 | spl24_215)), inference(forward_demodulation, [], [f1703, f695])).
fof(f695, plain, ((e0 = op(e4, e4)) | ~ spl24_1), inference(avatar_component_clause, [], [f693])).
fof(f1703, plain, (~ (e3 = op(e4, op(e4, e4))) | spl24_215), inference(avatar_component_clause, [], [f1701])).
fof(f12285, plain, (~ spl24_13 | ~ spl24_18), inference(avatar_split_clause, [], [f12284, f764, f743])).
fof(f12284, plain, (~ (e2 = op(e4, e2)) | ~ spl24_18), inference(forward_demodulation, [], [f476, f766])).
fof(f12283, plain, (~ spl24_13 | ~ spl24_38), inference(avatar_split_clause, [], [f12282, f848, f743])).
fof(f12282, plain, (~ (e2 = op(e4, e2)) | ~ spl24_38), inference(forward_demodulation, [], [f411, f850])).
fof(f12281, plain, (~ spl24_15 | ~ spl24_10), inference(avatar_split_clause, [], [f12280, f730, f751])).
fof(f12280, plain, (~ (e4 = op(e4, e2)) | ~ spl24_10), inference(forward_demodulation, [], [f479, f732])).
fof(f732, plain, ((e4 = op(e4, e3)) | ~ spl24_10), inference(avatar_component_clause, [], [f730])).
fof(f479, plain, ~ (op(e4, e2) = op(e4, e3)), inference(cnf_transformation, [], [f4])).
fof(f12245, plain, (~ spl24_103 | ~ spl24_123), inference(avatar_split_clause, [], [f12237, f1205, f1121])).
fof(f12237, plain, (~ (e2 = op(e0, e4)) | ~ spl24_123), inference(backward_demodulation, [], [f435, f1207])).
fof(f12235, plain, (spl24_78 | ~ spl24_12 | ~ spl24_136), inference(avatar_split_clause, [], [f12234, f1312, f739, f1016])).
fof(f12234, plain, ((e2 = op(e1, e4)) | (~ spl24_12 | ~ spl24_136)), inference(forward_demodulation, [], [f1313, f741])).
fof(f1313, plain, ((e2 = op(op(e4, e2), e4)) | ~ spl24_136), inference(avatar_component_clause, [], [f1312])).
fof(f12222, plain, (~ spl24_100 | ~ spl24_91 | spl24_229), inference(avatar_split_clause, [], [f12221, f1764, f1071, f1108])).
fof(f12221, plain, (~ (e4 = op(e1, e0)) | (~ spl24_91 | spl24_229)), inference(forward_demodulation, [], [f1766, f1073])).
fof(f1766, plain, (~ (e4 = op(e1, op(e1, e1))) | spl24_229), inference(avatar_component_clause, [], [f1764])).
fof(f12136, plain, (~ spl24_46 | ~ spl24_106 | spl24_163), inference(avatar_contradiction_clause, [], [f12135])).
fof(f12135, plain, ($false | (~ spl24_46 | ~ spl24_106 | spl24_163)), inference(subsumption_resolution, [], [f12134, f1136])).
fof(f12134, plain, (~ (e0 = op(e0, e3)) | (~ spl24_46 | spl24_163)), inference(forward_demodulation, [], [f1445, f884])).
fof(f884, plain, ((e0 = op(e3, e0)) | ~ spl24_46), inference(avatar_component_clause, [], [f882])).
fof(f12113, plain, (~ spl24_85 | ~ spl24_10), inference(avatar_split_clause, [], [f12112, f730, f1045])).
fof(f12112, plain, (~ (e4 = op(e1, e3)) | ~ spl24_10), inference(forward_demodulation, [], [f418, f732])).
fof(f12110, plain, (~ spl24_82 | ~ spl24_42 | spl24_159), inference(avatar_split_clause, [], [f12109, f1424, f865, f1033])).
fof(f12109, plain, (~ (e1 = op(e1, e3)) | (~ spl24_42 | spl24_159)), inference(forward_demodulation, [], [f1426, f867])).
fof(f12097, plain, (~ spl24_38 | ~ spl24_58 | spl24_153), inference(avatar_contradiction_clause, [], [f12096])).
fof(f12096, plain, ($false | (~ spl24_38 | ~ spl24_58 | spl24_153)), inference(subsumption_resolution, [], [f12095, f934])).
fof(f12095, plain, (~ (e2 = op(e2, e3)) | (~ spl24_38 | spl24_153)), inference(forward_demodulation, [], [f1397, f850])).
fof(f12090, plain, (~ spl24_38 | ~ spl24_58 | spl24_155), inference(avatar_contradiction_clause, [], [f12089])).
fof(f12089, plain, ($false | (~ spl24_38 | ~ spl24_58 | spl24_155)), inference(subsumption_resolution, [], [f12088, f934])).
fof(f12088, plain, (~ (e2 = op(e2, e3)) | (~ spl24_38 | spl24_155)), inference(forward_demodulation, [], [f1407, f850])).
fof(f12074, plain, (~ spl24_24 | ~ spl24_46 | ~ spl24_145), inference(avatar_contradiction_clause, [], [f12073])).
fof(f12073, plain, ($false | (~ spl24_24 | ~ spl24_46 | ~ spl24_145)), inference(subsumption_resolution, [], [f12072, f485])).
fof(f12072, plain, ((e0 = e4) | (~ spl24_24 | ~ spl24_46 | ~ spl24_145)), inference(forward_demodulation, [], [f12070, f884])).
fof(f12070, plain, ((e4 = op(e3, e0)) | (~ spl24_24 | ~ spl24_145)), inference(backward_demodulation, [], [f1357, f791])).
fof(f12036, plain, (~ spl24_27 | ~ spl24_42), inference(avatar_split_clause, [], [f12030, f865, f802])).
fof(f12030, plain, (~ (e1 = op(e3, e4)) | ~ spl24_42), inference(backward_demodulation, [], [f468, f867])).
fof(f12026, plain, (~ spl24_46 | ~ spl24_106 | spl24_161), inference(avatar_contradiction_clause, [], [f12025])).
fof(f12025, plain, ($false | (~ spl24_46 | ~ spl24_106 | spl24_161)), inference(subsumption_resolution, [], [f12022, f1136])).
fof(f12022, plain, (~ (e0 = op(e0, e3)) | (~ spl24_46 | spl24_161)), inference(backward_demodulation, [], [f1435, f884])).
fof(f12013, plain, (~ spl24_4 | ~ spl24_54), inference(avatar_split_clause, [], [f12009, f915, f705])).
fof(f12009, plain, (~ (e3 = op(e4, e4)) | ~ spl24_54), inference(backward_demodulation, [], [f430, f917])).
fof(f12004, plain, (~ spl24_33 | ~ spl24_58), inference(avatar_split_clause, [], [f11997, f932, f827])).
fof(f11997, plain, (~ (e2 = op(e3, e3)) | ~ spl24_58), inference(backward_demodulation, [], [f419, f934])).
fof(f11986, plain, (~ spl24_72 | ~ spl24_119 | ~ spl24_252 | spl24_265), inference(avatar_contradiction_clause, [], [f11985])).
fof(f11985, plain, ($false | (~ spl24_72 | ~ spl24_119 | ~ spl24_252 | spl24_265)), inference(subsumption_resolution, [], [f11982, f1190])).
fof(f11982, plain, (~ (e3 = op(e0, e1)) | (~ spl24_72 | ~ spl24_252 | spl24_265)), inference(backward_demodulation, [], [f11815, f993])).
fof(f11815, plain, (~ (e3 = op(e0, op(e2, e0))) | (~ spl24_252 | spl24_265)), inference(backward_demodulation, [], [f1982, f1889])).
fof(f11934, plain, (spl24_72 | ~ spl24_98 | ~ spl24_193), inference(avatar_split_clause, [], [f11931, f1592, f1100, f991])).
fof(f11931, plain, ((e1 = op(e2, e0)) | (~ spl24_98 | ~ spl24_193)), inference(backward_demodulation, [], [f1594, f1102])).
fof(f11890, plain, (~ spl24_105 | ~ spl24_115), inference(avatar_split_clause, [], [f11886, f1171, f1129])).
fof(f11886, plain, (~ (e4 = op(e0, e4)) | ~ spl24_115), inference(backward_demodulation, [], [f440, f1173])).
fof(f11881, plain, (~ spl24_109 | ~ spl24_119), inference(avatar_split_clause, [], [f11874, f1188, f1146])).
fof(f11874, plain, (~ (e3 = op(e0, e3)) | ~ spl24_119), inference(backward_demodulation, [], [f437, f1190])).
fof(f11820, plain, (~ spl24_86 | ~ spl24_237 | spl24_250), inference(avatar_split_clause, [], [f11819, f1876, f1805, f1050])).
fof(f1805, plain, (spl24_237 <=> (e1 = op(e2, op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl24_237])])).
fof(f11819, plain, (~ (e0 = op(e1, e2)) | (~ spl24_237 | spl24_250)), inference(forward_demodulation, [], [f1878, f1806])).
fof(f1806, plain, ((e1 = op(e2, op(e2, e2))) | ~ spl24_237), inference(avatar_component_clause, [], [f1805])).
fof(f11770, plain, (~ spl24_65 | ~ spl24_70), inference(avatar_split_clause, [], [f11769, f982, f961])).
fof(f11769, plain, (~ (e4 = op(e2, e2)) | ~ spl24_70), inference(forward_demodulation, [], [f456, f984])).
fof(f456, plain, ~ (op(e2, e1) = op(e2, e2)), inference(cnf_transformation, [], [f4])).
fof(f11760, plain, (~ spl24_60 | ~ spl24_40 | spl24_155), inference(avatar_split_clause, [], [f11759, f1405, f856, f940])).
fof(f11759, plain, (~ (e4 = op(e2, e3)) | (~ spl24_40 | spl24_155)), inference(forward_demodulation, [], [f1407, f858])).
fof(f11752, plain, (~ spl24_45 | ~ spl24_40), inference(avatar_split_clause, [], [f11523, f856, f877])).
fof(f11523, plain, (~ (e4 = op(e3, e1)) | ~ spl24_40), inference(forward_demodulation, [], [f466, f858])).
fof(f11735, plain, (~ spl24_7 | ~ spl24_12), inference(avatar_split_clause, [], [f11734, f739, f718])).
fof(f11734, plain, (~ (e1 = op(e4, e3)) | ~ spl24_12), inference(forward_demodulation, [], [f479, f741])).
fof(f11719, plain, (~ spl24_6 | ~ spl24_85 | ~ spl24_185), inference(avatar_contradiction_clause, [], [f11718])).
fof(f11718, plain, ($false | (~ spl24_6 | ~ spl24_85 | ~ spl24_185)), inference(subsumption_resolution, [], [f11717, f482])).
fof(f482, plain, ~ (e0 = e1), inference(cnf_transformation, [], [f5])).
fof(f11717, plain, ((e0 = e1) | (~ spl24_6 | ~ spl24_85 | ~ spl24_185)), inference(forward_demodulation, [], [f11716, f716])).
fof(f716, plain, ((e0 = op(e4, e3)) | ~ spl24_6), inference(avatar_component_clause, [], [f714])).
fof(f11716, plain, ((e1 = op(e4, e3)) | (~ spl24_85 | ~ spl24_185)), inference(forward_demodulation, [], [f1553, f1047])).
fof(f11706, plain, (~ spl24_18 | ~ spl24_85 | ~ spl24_94 | spl24_235), inference(avatar_contradiction_clause, [], [f11705])).
fof(f11705, plain, ($false | (~ spl24_18 | ~ spl24_85 | ~ spl24_94 | spl24_235)), inference(subsumption_resolution, [], [f11704, f766])).
fof(f11704, plain, (~ (e2 = op(e4, e1)) | (~ spl24_85 | ~ spl24_94 | spl24_235)), inference(forward_demodulation, [], [f11703, f1047])).
fof(f11703, plain, (~ (e2 = op(op(e1, e3), e1)) | (~ spl24_94 | spl24_235)), inference(forward_demodulation, [], [f1796, f1085])).
fof(f11652, plain, (spl24_18 | ~ spl24_78 | ~ spl24_142), inference(avatar_split_clause, [], [f11646, f1341, f1016, f764])).
fof(f11646, plain, ((e2 = op(e4, e1)) | (~ spl24_78 | ~ spl24_142)), inference(backward_demodulation, [], [f1342, f1018])).
fof(f11640, plain, (~ spl24_85 | ~ spl24_94 | spl24_229), inference(avatar_contradiction_clause, [], [f11639])).
fof(f11639, plain, ($false | (~ spl24_85 | ~ spl24_94 | spl24_229)), inference(subsumption_resolution, [], [f11633, f1047])).
fof(f11633, plain, (~ (e4 = op(e1, e3)) | (~ spl24_94 | spl24_229)), inference(backward_demodulation, [], [f1766, f1085])).
fof(f11615, plain, (~ spl24_86 | spl24_112 | ~ spl24_188), inference(avatar_contradiction_clause, [], [f11614])).
fof(f11614, plain, ($false | (~ spl24_86 | spl24_112 | ~ spl24_188)), inference(subsumption_resolution, [], [f11613, f1160])).
fof(f1160, plain, (~ (e1 = op(e0, e2)) | spl24_112), inference(avatar_component_clause, [], [f1159])).
fof(f11613, plain, ((e1 = op(e0, e2)) | (~ spl24_86 | ~ spl24_188)), inference(forward_demodulation, [], [f1568, f1052])).
fof(f11511, plain, (~ spl24_80 | ~ spl24_70 | ~ spl24_142), inference(avatar_split_clause, [], [f11510, f1341, f982, f1024])).
fof(f11510, plain, (~ (e4 = op(e1, e4)) | (~ spl24_70 | ~ spl24_142)), inference(forward_demodulation, [], [f11186, f984])).
fof(f11186, plain, (~ (op(e1, e4) = op(e2, e1)) | ~ spl24_142), inference(forward_demodulation, [], [f400, f1342])).
fof(f11509, plain, (~ spl24_78 | spl24_214 | ~ spl24_230), inference(avatar_split_clause, [], [f11254, f1769, f1697, f1016])).
fof(f1697, plain, (spl24_214 <=> (e2 = op(op(e4, op(e4, e4)), e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_214])])).
fof(f11254, plain, (~ (e2 = op(e1, e4)) | (spl24_214 | ~ spl24_230)), inference(backward_demodulation, [], [f1699, f1770])).
fof(f1699, plain, (~ (e2 = op(op(e4, op(e4, e4)), e4)) | spl24_214), inference(avatar_component_clause, [], [f1697])).
fof(f11508, plain, (spl24_18 | ~ spl24_70 | ~ spl24_174), inference(avatar_split_clause, [], [f11477, f1498, f982, f764])).
fof(f11477, plain, ((e2 = op(e4, e1)) | (~ spl24_70 | ~ spl24_174)), inference(backward_demodulation, [], [f1500, f984])).
fof(f11489, plain, (~ spl24_57 | spl24_83 | ~ spl24_169), inference(avatar_contradiction_clause, [], [f11488])).
fof(f11488, plain, ($false | (~ spl24_57 | spl24_83 | ~ spl24_169)), inference(subsumption_resolution, [], [f11485, f1038])).
fof(f1038, plain, (~ (e2 = op(e1, e3)) | spl24_83), inference(avatar_component_clause, [], [f1037])).
fof(f11485, plain, ((e2 = op(e1, e3)) | (~ spl24_57 | ~ spl24_169)), inference(backward_demodulation, [], [f1474, f930])).
fof(f11442, plain, (~ spl24_36 | ~ spl24_86), inference(avatar_split_clause, [], [f11433, f1050, f840])).
fof(f11433, plain, (~ (e0 = op(e3, e2)) | ~ spl24_86), inference(backward_demodulation, [], [f407, f1052])).
fof(f11432, plain, (~ spl24_86 | ~ spl24_93 | spl24_261), inference(avatar_split_clause, [], [f11428, f1949, f1079, f1050])).
fof(f11428, plain, (~ (e0 = op(e1, e2)) | (~ spl24_93 | spl24_261)), inference(backward_demodulation, [], [f1951, f1081])).
fof(f11430, plain, (~ spl24_88 | ~ spl24_93), inference(avatar_split_clause, [], [f11420, f1079, f1058])).
fof(f11420, plain, (~ (e2 = op(e1, e2)) | ~ spl24_93), inference(backward_demodulation, [], [f446, f1081])).
fof(f11429, plain, (~ spl24_68 | ~ spl24_93), inference(avatar_split_clause, [], [f11417, f1079, f974])).
fof(f11417, plain, (~ (e2 = op(e2, e1)) | ~ spl24_93), inference(backward_demodulation, [], [f396, f1081])).
fof(f11411, plain, (~ spl24_80 | ~ spl24_105), inference(avatar_split_clause, [], [f11406, f1129, f1024])).
fof(f11406, plain, (~ (e4 = op(e1, e4)) | ~ spl24_105), inference(backward_demodulation, [], [f422, f1131])).
fof(f11403, plain, (~ spl24_84 | ~ spl24_109), inference(avatar_split_clause, [], [f11395, f1146, f1041])).
fof(f11395, plain, (~ (e3 = op(e1, e3)) | ~ spl24_109), inference(backward_demodulation, [], [f412, f1148])).
fof(f11390, plain, (~ spl24_88 | ~ spl24_113), inference(avatar_split_clause, [], [f11378, f1163, f1058])).
fof(f11378, plain, (~ (e2 = op(e1, e2)) | ~ spl24_113), inference(backward_demodulation, [], [f402, f1165])).
fof(f11369, plain, (spl24_91 | ~ spl24_117 | ~ spl24_206), inference(avatar_contradiction_clause, [], [f11368])).
fof(f11368, plain, ($false | (spl24_91 | ~ spl24_117 | ~ spl24_206)), inference(subsumption_resolution, [], [f11354, f1072])).
fof(f1072, plain, (~ (e0 = op(e1, e1)) | spl24_91), inference(avatar_component_clause, [], [f1071])).
fof(f11354, plain, ((e0 = op(e1, e1)) | (~ spl24_117 | ~ spl24_206)), inference(backward_demodulation, [], [f1658, f1182])).
fof(f1658, plain, ((e0 = op(op(e0, e1), e1)) | ~ spl24_206), inference(avatar_component_clause, [], [f1656])).
fof(f11342, plain, (~ spl24_121 | ~ spl24_260), inference(avatar_contradiction_clause, [], [f11341])).
fof(f11341, plain, ($false | (~ spl24_121 | ~ spl24_260)), inference(subsumption_resolution, [], [f11340, f482])).
fof(f11340, plain, ((e0 = e1) | (~ spl24_121 | ~ spl24_260)), inference(forward_demodulation, [], [f11337, f1199])).
fof(f11337, plain, ((op(e0, e0) = e1) | (~ spl24_121 | ~ spl24_260)), inference(backward_demodulation, [], [f1945, f1199])).
fof(f11339, plain, (~ spl24_116 | ~ spl24_121), inference(avatar_split_clause, [], [f11329, f1197, f1176])).
fof(f11329, plain, (~ (e0 = op(e0, e1)) | ~ spl24_121), inference(backward_demodulation, [], [f432, f1199])).
fof(f11338, plain, (~ spl24_96 | ~ spl24_121), inference(avatar_split_clause, [], [f11325, f1197, f1092])).
fof(f11325, plain, (~ (e0 = op(e1, e0)) | ~ spl24_121), inference(backward_demodulation, [], [f382, f1199])).
fof(f11295, plain, (spl24_79 | ~ spl24_27 | ~ spl24_149), inference(avatar_split_clause, [], [f11294, f1374, f802, f1020])).
fof(f11294, plain, ((e3 = op(e1, e4)) | (~ spl24_27 | ~ spl24_149)), inference(forward_demodulation, [], [f1376, f804])).
fof(f804, plain, ((e1 = op(e3, e4)) | ~ spl24_27), inference(avatar_component_clause, [], [f802])).
fof(f11285, plain, (~ spl24_115 | ~ spl24_51 | spl24_165), inference(avatar_split_clause, [], [f11284, f1452, f903, f1171])).
fof(f11284, plain, (~ (e4 = op(e0, e2)) | (~ spl24_51 | spl24_165)), inference(forward_demodulation, [], [f1454, f905])).
fof(f11251, plain, (~ spl24_86 | spl24_227 | ~ spl24_235), inference(avatar_split_clause, [], [f11245, f1794, f1756, f1050])).
fof(f1756, plain, (spl24_227 <=> (e0 = op(e1, op(op(e1, op(e1, e1)), e1)))), introduced(avatar_definition, [new_symbols(naming, [spl24_227])])).
fof(f11245, plain, (~ (e0 = op(e1, e2)) | (spl24_227 | ~ spl24_235)), inference(backward_demodulation, [], [f1758, f1795])).
fof(f1795, plain, ((e2 = op(op(e1, op(e1, e1)), e1)) | ~ spl24_235), inference(avatar_component_clause, [], [f1794])).
fof(f1758, plain, (~ (e0 = op(e1, op(op(e1, op(e1, e1)), e1))) | spl24_227), inference(avatar_component_clause, [], [f1756])).
fof(f11238, plain, (~ spl24_36 | ~ spl24_33 | spl24_249), inference(avatar_split_clause, [], [f11237, f1871, f827, f840])).
fof(f11237, plain, (~ (e0 = op(e3, e2)) | (~ spl24_33 | spl24_249)), inference(forward_demodulation, [], [f1873, f829])).
fof(f11188, plain, (~ spl24_69 | ~ spl24_64), inference(avatar_split_clause, [], [f11187, f957, f978])).
fof(f11187, plain, (~ (e3 = op(e2, e1)) | ~ spl24_64), inference(forward_demodulation, [], [f456, f959])).
fof(f11185, plain, (~ spl24_69 | ~ spl24_234 | spl24_262), inference(avatar_split_clause, [], [f11001, f1960, f1788, f978])).
fof(f11001, plain, (~ (e3 = op(e2, e1)) | (~ spl24_234 | spl24_262)), inference(forward_demodulation, [], [f1962, f1789])).
fof(f1789, plain, ((e1 = op(op(e2, op(e2, e2)), e2)) | ~ spl24_234), inference(avatar_component_clause, [], [f1788])).
fof(f11167, plain, (~ spl24_37 | ~ spl24_12), inference(avatar_split_clause, [], [f11166, f739, f844])).
fof(f11166, plain, (~ (e1 = op(e3, e2)) | ~ spl24_12), inference(forward_demodulation, [], [f411, f741])).
fof(f11165, plain, (~ spl24_38 | ~ spl24_33), inference(avatar_split_clause, [], [f11164, f827, f848])).
fof(f11164, plain, (~ (e2 = op(e3, e2)) | ~ spl24_33), inference(forward_demodulation, [], [f469, f829])).
fof(f11154, plain, (~ spl24_12 | ~ spl24_88 | ~ spl24_137), inference(avatar_contradiction_clause, [], [f11153])).
fof(f11153, plain, ($false | (~ spl24_12 | ~ spl24_88 | ~ spl24_137)), inference(subsumption_resolution, [], [f11152, f490])).
fof(f11152, plain, ((e2 = e4) | (~ spl24_12 | ~ spl24_88 | ~ spl24_137)), inference(forward_demodulation, [], [f11151, f1060])).
fof(f11151, plain, ((e4 = op(e1, e2)) | (~ spl24_12 | ~ spl24_137)), inference(forward_demodulation, [], [f1319, f741])).
fof(f11079, plain, (~ spl24_72 | ~ spl24_96 | ~ spl24_178), inference(avatar_contradiction_clause, [], [f11078])).
fof(f11078, plain, ($false | (~ spl24_72 | ~ spl24_96 | ~ spl24_178)), inference(subsumption_resolution, [], [f11077, f483])).
fof(f483, plain, ~ (e0 = e2), inference(cnf_transformation, [], [f5])).
fof(f11077, plain, ((e0 = e2) | (~ spl24_72 | ~ spl24_96 | ~ spl24_178)), inference(forward_demodulation, [], [f11075, f1094])).
fof(f1094, plain, ((e0 = op(e1, e0)) | ~ spl24_96), inference(avatar_component_clause, [], [f1092])).
fof(f11075, plain, ((e2 = op(e1, e0)) | (~ spl24_72 | ~ spl24_178)), inference(backward_demodulation, [], [f1519, f993])).
fof(f11048, plain, (~ spl24_104 | ~ spl24_6 | spl24_132), inference(avatar_split_clause, [], [f11047, f1293, f714, f1125])).
fof(f11047, plain, (~ (e3 = op(e0, e4)) | (~ spl24_6 | spl24_132)), inference(forward_demodulation, [], [f1295, f716])).
fof(f11023, plain, (~ spl24_245 | ~ spl24_260), inference(avatar_contradiction_clause, [], [f11022])).
fof(f11022, plain, ($false | (~ spl24_245 | ~ spl24_260)), inference(subsumption_resolution, [], [f11020, f488])).
fof(f11020, plain, ((e1 = e4) | (~ spl24_245 | ~ spl24_260)), inference(backward_demodulation, [], [f1853, f1945])).
fof(f10991, plain, (~ spl24_22 | ~ spl24_96 | ~ spl24_145), inference(avatar_contradiction_clause, [], [f10990])).
fof(f10990, plain, ($false | (~ spl24_22 | ~ spl24_96 | ~ spl24_145)), inference(subsumption_resolution, [], [f10989, f485])).
fof(f10989, plain, ((e0 = e4) | (~ spl24_22 | ~ spl24_96 | ~ spl24_145)), inference(forward_demodulation, [], [f10986, f1094])).
fof(f10986, plain, ((e4 = op(e1, e0)) | (~ spl24_22 | ~ spl24_145)), inference(backward_demodulation, [], [f1357, f783])).
fof(f10956, plain, (~ spl24_23 | ~ spl24_123), inference(avatar_split_clause, [], [f10949, f1205, f785])).
fof(f10949, plain, (~ (e2 = op(e4, e0)) | ~ spl24_123), inference(backward_demodulation, [], [f385, f1207])).
fof(f10946, plain, (~ spl24_108 | ~ spl24_36 | spl24_153), inference(avatar_split_clause, [], [f10945, f1395, f840, f1142])).
fof(f10945, plain, (~ (e2 = op(e0, e3)) | (~ spl24_36 | spl24_153)), inference(forward_demodulation, [], [f1397, f842])).
fof(f842, plain, ((e0 = op(e3, e2)) | ~ spl24_36), inference(avatar_component_clause, [], [f840])).
fof(f10938, plain, (~ spl24_23 | ~ spl24_115 | spl24_202), inference(avatar_split_clause, [], [f10937, f1636, f1171, f785])).
fof(f10937, plain, (~ (e2 = op(e4, e0)) | (~ spl24_115 | spl24_202)), inference(forward_demodulation, [], [f1638, f1173])).
fof(f10898, plain, (~ spl24_110 | ~ spl24_6 | spl24_133), inference(avatar_split_clause, [], [f10860, f1298, f714, f1150])).
fof(f10860, plain, (~ (e4 = op(e0, e3)) | (~ spl24_6 | spl24_133)), inference(backward_demodulation, [], [f1299, f716])).
fof(f1299, plain, (~ (e4 = op(op(e4, e3), e3)) | spl24_133), inference(avatar_component_clause, [], [f1298])).
fof(f10895, plain, (~ spl24_30 | ~ spl24_50), inference(avatar_split_clause, [], [f10894, f898, f814])).
fof(f10894, plain, (~ (e4 = op(e3, e4)) | ~ spl24_50), inference(forward_demodulation, [], [f465, f900])).
fof(f10890, plain, (~ spl24_25 | ~ spl24_50), inference(avatar_split_clause, [], [f10889, f898, f793])).
fof(f10889, plain, (~ (e4 = op(e4, e0)) | ~ spl24_50), inference(forward_demodulation, [], [f391, f900])).
fof(f391, plain, ~ (op(e3, e0) = op(e4, e0)), inference(cnf_transformation, [], [f4])).
fof(f10888, plain, (~ spl24_24 | ~ spl24_74), inference(avatar_split_clause, [], [f10887, f999, f789])).
fof(f10887, plain, (~ (e3 = op(e4, e0)) | ~ spl24_74), inference(forward_demodulation, [], [f390, f1001])).
fof(f10884, plain, (~ spl24_24 | ~ spl24_14), inference(avatar_split_clause, [], [f10482, f747, f789])).
fof(f10482, plain, (~ (e3 = op(e4, e0)) | ~ spl24_14), inference(forward_demodulation, [], [f473, f749])).
fof(f10861, plain, (~ spl24_1 | ~ spl24_6), inference(avatar_split_clause, [], [f10854, f714, f693])).
fof(f10854, plain, (~ (e0 = op(e4, e4)) | ~ spl24_6), inference(backward_demodulation, [], [f481, f716])).
fof(f10846, plain, (~ spl24_42 | ~ spl24_44), inference(avatar_contradiction_clause, [], [f10845])).
fof(f10845, plain, ($false | (~ spl24_42 | ~ spl24_44)), inference(subsumption_resolution, [], [f10844, f487])).
fof(f10844, plain, ((e1 = e3) | (~ spl24_42 | ~ spl24_44)), inference(backward_demodulation, [], [f875, f867])).
fof(f10840, plain, (~ spl24_35 | ~ spl24_50), inference(avatar_split_clause, [], [f10837, f898, f835])).
fof(f10837, plain, (~ (e4 = op(e3, e3)) | ~ spl24_50), inference(backward_demodulation, [], [f464, f900])).
fof(f10827, plain, (~ spl24_35 | ~ spl24_60), inference(avatar_split_clause, [], [f10821, f940, f835])).
fof(f10821, plain, (~ (e4 = op(e3, e3)) | ~ spl24_60), inference(backward_demodulation, [], [f419, f942])).
fof(f10818, plain, (~ spl24_54 | ~ spl24_74), inference(avatar_split_clause, [], [f10812, f999, f915])).
fof(f10812, plain, (~ (e3 = op(e2, e4)) | ~ spl24_74), inference(backward_demodulation, [], [f455, f1001])).
fof(f10809, plain, (~ spl24_82 | ~ spl24_84), inference(avatar_contradiction_clause, [], [f10808])).
fof(f10808, plain, ($false | (~ spl24_82 | ~ spl24_84)), inference(subsumption_resolution, [], [f10807, f487])).
fof(f10807, plain, ((e1 = e3) | (~ spl24_82 | ~ spl24_84)), inference(backward_demodulation, [], [f1043, f1035])).
fof(f10805, plain, (~ spl24_50 | ~ spl24_104 | spl24_196), inference(avatar_split_clause, [], [f10802, f1606, f1125, f898])).
fof(f10802, plain, (~ (e4 = op(e3, e0)) | (~ spl24_104 | spl24_196)), inference(backward_demodulation, [], [f1608, f1127])).
fof(f10793, plain, (~ spl24_47 | ~ spl24_122), inference(avatar_split_clause, [], [f10786, f1201, f886])).
fof(f10786, plain, (~ (e1 = op(e3, e0)) | ~ spl24_122), inference(backward_demodulation, [], [f384, f1203])).
fof(f10776, plain, (spl24_60 | ~ spl24_28 | ~ spl24_148), inference(avatar_split_clause, [], [f10775, f1369, f806, f940])).
fof(f10775, plain, ((e4 = op(e2, e3)) | (~ spl24_28 | ~ spl24_148)), inference(forward_demodulation, [], [f1370, f808])).
fof(f1370, plain, ((e4 = op(op(e3, e4), e3)) | ~ spl24_148), inference(avatar_component_clause, [], [f1369])).
fof(f10772, plain, (~ spl24_56 | ~ spl24_36 | spl24_155), inference(avatar_split_clause, [], [f10771, f1405, f840, f924])).
fof(f10771, plain, (~ (e0 = op(e2, e3)) | (~ spl24_36 | spl24_155)), inference(forward_demodulation, [], [f1407, f842])).
fof(f10770, plain, (~ spl24_44 | ~ spl24_84 | spl24_159), inference(avatar_contradiction_clause, [], [f10769])).
fof(f10769, plain, ($false | (~ spl24_44 | ~ spl24_84 | spl24_159)), inference(subsumption_resolution, [], [f10768, f1043])).
fof(f10768, plain, (~ (e3 = op(e1, e3)) | (~ spl24_44 | spl24_159)), inference(forward_demodulation, [], [f1426, f875])).
fof(f10761, plain, (spl24_74 | ~ spl24_108 | ~ spl24_199), inference(avatar_split_clause, [], [f10760, f1621, f1142, f999])).
fof(f10760, plain, ((e3 = op(e2, e0)) | (~ spl24_108 | ~ spl24_199)), inference(forward_demodulation, [], [f1622, f1144])).
fof(f1622, plain, ((e3 = op(op(e0, e3), e0)) | ~ spl24_199), inference(avatar_component_clause, [], [f1621])).
fof(f10739, plain, (~ spl24_57 | ~ spl24_224 | spl24_225), inference(avatar_split_clause, [], [f10738, f1746, f1741, f928])).
fof(f10738, plain, (~ (e1 = op(e2, e3)) | (~ spl24_224 | spl24_225)), inference(forward_demodulation, [], [f1748, f1742])).
fof(f10692, plain, (~ spl24_56 | ~ spl24_108 | spl24_169), inference(avatar_contradiction_clause, [], [f10691])).
fof(f10691, plain, ($false | (~ spl24_56 | ~ spl24_108 | spl24_169)), inference(subsumption_resolution, [], [f10687, f1144])).
fof(f10687, plain, (~ (e2 = op(e0, e3)) | (~ spl24_56 | spl24_169)), inference(backward_demodulation, [], [f1473, f926])).
fof(f1473, plain, (~ (e2 = op(op(e2, e3), e3)) | spl24_169), inference(avatar_component_clause, [], [f1472])).
fof(f10690, plain, (~ spl24_51 | ~ spl24_56), inference(avatar_split_clause, [], [f10681, f924, f903])).
fof(f10681, plain, (~ (e0 = op(e2, e4)) | ~ spl24_56), inference(backward_demodulation, [], [f461, f926])).
fof(f10669, plain, (spl24_56 | ~ spl24_108 | ~ spl24_200), inference(avatar_split_clause, [], [f10667, f1626, f1142, f924])).
fof(f10667, plain, ((e0 = op(e2, e3)) | (~ spl24_108 | ~ spl24_200)), inference(backward_demodulation, [], [f1628, f1144])).
fof(f10668, plain, (~ spl24_58 | ~ spl24_108), inference(avatar_split_clause, [], [f10661, f1142, f932])).
fof(f10661, plain, (~ (e2 = op(e2, e3)) | ~ spl24_108), inference(backward_demodulation, [], [f413, f1144])).
fof(f10648, plain, (spl24_108 | ~ spl24_36 | ~ spl24_153), inference(avatar_split_clause, [], [f10647, f1395, f840, f1142])).
fof(f10647, plain, ((e2 = op(e0, e3)) | (~ spl24_36 | ~ spl24_153)), inference(forward_demodulation, [], [f1396, f842])).
fof(f1396, plain, ((e2 = op(op(e3, e2), e3)) | ~ spl24_153), inference(avatar_component_clause, [], [f1395])).
fof(f10632, plain, (~ spl24_104 | ~ spl24_75 | ~ spl24_252 | spl24_265), inference(avatar_split_clause, [], [f10631, f1980, f1888, f1003, f1125])).
fof(f10631, plain, (~ (e3 = op(e0, e4)) | (~ spl24_75 | ~ spl24_252 | spl24_265)), inference(forward_demodulation, [], [f10626, f1005])).
fof(f10626, plain, (~ (e3 = op(e0, op(e2, e0))) | (~ spl24_252 | spl24_265)), inference(backward_demodulation, [], [f1982, f1889])).
fof(f10628, plain, (~ spl24_102 | ~ spl24_75 | spl24_243 | ~ spl24_252), inference(avatar_split_clause, [], [f10627, f1888, f1844, f1003, f1117])).
fof(f10627, plain, (~ (e1 = op(e0, e4)) | (~ spl24_75 | spl24_243 | ~ spl24_252)), inference(forward_demodulation, [], [f10620, f1005])).
fof(f10620, plain, (~ (e1 = op(e0, op(e2, e0))) | (spl24_243 | ~ spl24_252)), inference(backward_demodulation, [], [f1846, f1889])).
fof(f10609, plain, (~ spl24_47 | spl24_81 | ~ spl24_161), inference(avatar_contradiction_clause, [], [f10608])).
fof(f10608, plain, ($false | (~ spl24_47 | spl24_81 | ~ spl24_161)), inference(subsumption_resolution, [], [f10607, f1030])).
fof(f1030, plain, (~ (e0 = op(e1, e3)) | spl24_81), inference(avatar_component_clause, [], [f1029])).
fof(f1029, plain, (spl24_81 <=> (e0 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_81])])).
fof(f10607, plain, ((e0 = op(e1, e3)) | (~ spl24_47 | ~ spl24_161)), inference(forward_demodulation, [], [f1434, f888])).
fof(f1434, plain, ((e0 = op(op(e3, e0), e3)) | ~ spl24_161), inference(avatar_component_clause, [], [f1433])).
fof(f10601, plain, (~ spl24_56 | ~ spl24_28 | ~ spl24_35 | spl24_241), inference(avatar_split_clause, [], [f10589, f1834, f835, f806, f924])).
fof(f10589, plain, (~ (e0 = op(e2, e3)) | (~ spl24_28 | ~ spl24_35 | spl24_241)), inference(forward_demodulation, [], [f10582, f808])).
fof(f10582, plain, (~ (e0 = op(op(e3, e4), e3)) | (~ spl24_35 | spl24_241)), inference(backward_demodulation, [], [f1836, f837])).
fof(f1836, plain, (~ (e0 = op(op(e3, op(e3, e3)), e3)) | spl24_241), inference(avatar_component_clause, [], [f1834])).
fof(f10599, plain, (~ spl24_55 | ~ spl24_75), inference(avatar_split_clause, [], [f10598, f1003, f919])).
fof(f10598, plain, (~ (e4 = op(e2, e4)) | ~ spl24_75), inference(forward_demodulation, [], [f455, f1005])).
fof(f10596, plain, (spl24_54 | ~ spl24_28 | ~ spl24_149), inference(avatar_split_clause, [], [f10505, f1374, f806, f915])).
fof(f10505, plain, ((e3 = op(e2, e4)) | (~ spl24_28 | ~ spl24_149)), inference(forward_demodulation, [], [f1376, f808])).
fof(f10593, plain, (~ spl24_8 | ~ spl24_23), inference(avatar_split_clause, [], [f10592, f785, f722])).
fof(f10592, plain, (~ (e2 = op(e4, e3)) | ~ spl24_23), inference(forward_demodulation, [], [f474, f787])).
fof(f10585, plain, (~ spl24_28 | ~ spl24_35 | spl24_224), inference(avatar_contradiction_clause, [], [f10584])).
fof(f10584, plain, ($false | (~ spl24_28 | ~ spl24_35 | spl24_224)), inference(subsumption_resolution, [], [f10579, f808])).
fof(f10579, plain, (~ (e2 = op(e3, e4)) | (~ spl24_35 | spl24_224)), inference(backward_demodulation, [], [f1743, f837])).
fof(f10575, plain, (~ spl24_58 | ~ spl24_59), inference(avatar_contradiction_clause, [], [f10574])).
fof(f10574, plain, ($false | (~ spl24_58 | ~ spl24_59)), inference(subsumption_resolution, [], [f10573, f489])).
fof(f489, plain, ~ (e2 = e3), inference(cnf_transformation, [], [f5])).
fof(f10573, plain, ((e2 = e3) | (~ spl24_58 | ~ spl24_59)), inference(backward_demodulation, [], [f938, f934])).
fof(f10572, plain, (~ spl24_36 | ~ spl24_59 | ~ spl24_168), inference(avatar_contradiction_clause, [], [f10571])).
fof(f10571, plain, ($false | (~ spl24_36 | ~ spl24_59 | ~ spl24_168)), inference(subsumption_resolution, [], [f10570, f484])).
fof(f484, plain, ~ (e0 = e3), inference(cnf_transformation, [], [f5])).
fof(f10570, plain, ((e0 = e3) | (~ spl24_36 | ~ spl24_59 | ~ spl24_168)), inference(forward_demodulation, [], [f10568, f842])).
fof(f10568, plain, ((e3 = op(e3, e2)) | (~ spl24_59 | ~ spl24_168)), inference(backward_demodulation, [], [f1468, f938])).
fof(f10559, plain, (~ spl24_96 | ~ spl24_99), inference(avatar_contradiction_clause, [], [f10558])).
fof(f10558, plain, ($false | (~ spl24_96 | ~ spl24_99)), inference(subsumption_resolution, [], [f10557, f484])).
fof(f10557, plain, ((e0 = e3) | (~ spl24_96 | ~ spl24_99)), inference(forward_demodulation, [], [f1106, f1094])).
fof(f10556, plain, (spl24_81 | ~ spl24_107 | ~ spl24_200), inference(avatar_contradiction_clause, [], [f10555])).
fof(f10555, plain, ($false | (spl24_81 | ~ spl24_107 | ~ spl24_200)), inference(subsumption_resolution, [], [f10548, f1030])).
fof(f10548, plain, ((e0 = op(e1, e3)) | (~ spl24_107 | ~ spl24_200)), inference(backward_demodulation, [], [f1628, f1140])).
fof(f10474, plain, (spl24_75 | ~ spl24_23 | ~ spl24_145), inference(avatar_split_clause, [], [f10473, f1355, f785, f1003])).
fof(f10473, plain, ((e4 = op(e2, e0)) | (~ spl24_23 | ~ spl24_145)), inference(forward_demodulation, [], [f1357, f787])).
fof(f10454, plain, (spl24_50 | ~ spl24_104 | ~ spl24_196), inference(avatar_split_clause, [], [f10327, f1606, f1125, f898])).
fof(f10327, plain, ((e4 = op(e3, e0)) | (~ spl24_104 | ~ spl24_196)), inference(backward_demodulation, [], [f1607, f1127])).
fof(f1607, plain, ((e4 = op(op(e0, e4), e0)) | ~ spl24_196), inference(avatar_component_clause, [], [f1606])).
fof(f10362, plain, (~ spl24_12 | ~ spl24_62), inference(avatar_split_clause, [], [f10348, f949, f739])).
fof(f10348, plain, (~ (e1 = op(e4, e2)) | ~ spl24_62), inference(backward_demodulation, [], [f410, f951])).
fof(f10321, plain, (~ spl24_103 | ~ spl24_108), inference(avatar_split_clause, [], [f10315, f1142, f1121])).
fof(f10315, plain, (~ (e2 = op(e0, e4)) | ~ spl24_108), inference(backward_demodulation, [], [f441, f1144])).
fof(f10320, plain, (~ spl24_33 | ~ spl24_108), inference(avatar_split_clause, [], [f10313, f1142, f827])).
fof(f10313, plain, (~ (e2 = op(e3, e3)) | ~ spl24_108), inference(backward_demodulation, [], [f414, f1144])).
fof(f10308, plain, (~ spl24_107 | ~ spl24_122), inference(avatar_split_clause, [], [f10301, f1201, f1138])).
fof(f10301, plain, (~ (e1 = op(e0, e3)) | ~ spl24_122), inference(backward_demodulation, [], [f434, f1203])).
fof(f10253, plain, (~ spl24_63 | ~ spl24_88), inference(avatar_split_clause, [], [f10252, f1058, f953])).
fof(f10252, plain, (~ (e2 = op(e2, e2)) | ~ spl24_88), inference(forward_demodulation, [], [f406, f1060])).
fof(f10245, plain, (~ spl24_68 | ~ spl24_88 | spl24_175), inference(avatar_contradiction_clause, [], [f10244])).
fof(f10244, plain, ($false | (~ spl24_68 | ~ spl24_88 | spl24_175)), inference(subsumption_resolution, [], [f10243, f1060])).
fof(f10243, plain, (~ (e2 = op(e1, e2)) | (~ spl24_68 | spl24_175)), inference(forward_demodulation, [], [f1505, f976])).
fof(f10237, plain, (~ spl24_63 | ~ spl24_68), inference(avatar_split_clause, [], [f10236, f974, f953])).
fof(f10236, plain, (~ (e2 = op(e2, e2)) | ~ spl24_68), inference(forward_demodulation, [], [f456, f976])).
fof(f10233, plain, (~ spl24_68 | spl24_174), inference(avatar_contradiction_clause, [], [f10232])).
fof(f10232, plain, ($false | (~ spl24_68 | spl24_174)), inference(subsumption_resolution, [], [f10231, f976])).
fof(f10231, plain, (~ (e2 = op(e2, e1)) | (~ spl24_68 | spl24_174)), inference(forward_demodulation, [], [f1499, f976])).
fof(f1499, plain, (~ (e2 = op(op(e2, e1), e1)) | spl24_174), inference(avatar_component_clause, [], [f1498])).
fof(f10229, plain, (~ spl24_20 | ~ spl24_80 | spl24_142), inference(avatar_split_clause, [], [f10098, f1341, f1024, f772])).
fof(f10098, plain, (~ (e4 = op(e4, e1)) | (~ spl24_80 | spl24_142)), inference(backward_demodulation, [], [f1343, f1026])).
fof(f10210, plain, (~ spl24_10 | spl24_133), inference(avatar_contradiction_clause, [], [f10209])).
fof(f10209, plain, ($false | (~ spl24_10 | spl24_133)), inference(subsumption_resolution, [], [f10208, f732])).
fof(f10208, plain, (~ (e4 = op(e4, e3)) | (~ spl24_10 | spl24_133)), inference(backward_demodulation, [], [f1299, f732])).
fof(f10204, plain, (~ spl24_2 | ~ spl24_12), inference(avatar_split_clause, [], [f10199, f739, f697])).
fof(f10199, plain, (~ (e1 = op(e4, e4)) | ~ spl24_12), inference(backward_demodulation, [], [f480, f741])).
fof(f10164, plain, (spl24_24 | ~ spl24_50 | ~ spl24_162), inference(avatar_contradiction_clause, [], [f10163])).
fof(f10163, plain, ($false | (spl24_24 | ~ spl24_50 | ~ spl24_162)), inference(subsumption_resolution, [], [f10160, f790])).
fof(f790, plain, (~ (e3 = op(e4, e0)) | spl24_24), inference(avatar_component_clause, [], [f789])).
fof(f10160, plain, ((e3 = op(e4, e0)) | (~ spl24_50 | ~ spl24_162)), inference(backward_demodulation, [], [f1440, f900])).
fof(f1440, plain, ((e3 = op(op(e3, e0), e0)) | ~ spl24_162), inference(avatar_component_clause, [], [f1438])).
fof(f10118, plain, (~ spl24_68 | ~ spl24_70), inference(avatar_contradiction_clause, [], [f10117])).
fof(f10117, plain, ($false | (~ spl24_68 | ~ spl24_70)), inference(subsumption_resolution, [], [f10116, f490])).
fof(f10116, plain, ((e2 = e4) | (~ spl24_68 | ~ spl24_70)), inference(backward_demodulation, [], [f984, f976])).
fof(f10085, plain, (~ spl24_88 | ~ spl24_89), inference(avatar_contradiction_clause, [], [f10084])).
fof(f10084, plain, ($false | (~ spl24_88 | ~ spl24_89)), inference(subsumption_resolution, [], [f10083, f489])).
fof(f10083, plain, ((e2 = e3) | (~ spl24_88 | ~ spl24_89)), inference(backward_demodulation, [], [f1064, f1060])).
fof(f10079, plain, (~ spl24_84 | ~ spl24_89), inference(avatar_split_clause, [], [f10072, f1062, f1041])).
fof(f10072, plain, (~ (e3 = op(e1, e3)) | ~ spl24_89), inference(backward_demodulation, [], [f449, f1064])).
fof(f10029, plain, (~ spl24_96 | ~ spl24_116 | spl24_192), inference(avatar_contradiction_clause, [], [f10028])).
fof(f10028, plain, ($false | (~ spl24_96 | ~ spl24_116 | spl24_192)), inference(subsumption_resolution, [], [f10022, f1178])).
fof(f10022, plain, (~ (e0 = op(e0, e1)) | (~ spl24_96 | spl24_192)), inference(backward_demodulation, [], [f1589, f1094])).
fof(f10012, plain, (~ spl24_18 | ~ spl24_103 | ~ spl24_146), inference(avatar_split_clause, [], [f10004, f1360, f1121, f764])).
fof(f10004, plain, (~ (e2 = op(e4, e1)) | (~ spl24_103 | ~ spl24_146)), inference(backward_demodulation, [], [f9857, f1123])).
fof(f9857, plain, (~ (op(e0, e4) = op(e4, e1)) | ~ spl24_146), inference(forward_demodulation, [], [f472, f1361])).
fof(f10007, plain, (~ spl24_28 | ~ spl24_103), inference(avatar_split_clause, [], [f9995, f1121, f806])).
fof(f9995, plain, (~ (e2 = op(e3, e4)) | ~ spl24_103), inference(backward_demodulation, [], [f424, f1123])).
fof(f9971, plain, (~ spl24_90 | ~ spl24_115), inference(avatar_split_clause, [], [f9958, f1171, f1066])).
fof(f9958, plain, (~ (e4 = op(e1, e2)) | ~ spl24_115), inference(backward_demodulation, [], [f402, f1173])).
fof(f9955, plain, (~ spl24_96 | ~ spl24_116 | spl24_194), inference(avatar_split_clause, [], [f9952, f1597, f1176, f1092])).
fof(f9952, plain, (~ (e0 = op(e1, e0)) | (~ spl24_116 | spl24_194)), inference(backward_demodulation, [], [f1599, f1178])).
fof(f9940, plain, (spl24_107 | ~ spl24_124 | ~ spl24_260), inference(avatar_split_clause, [], [f9937, f1944, f1209, f1138])).
fof(f9937, plain, ((e1 = op(e0, e3)) | (~ spl24_124 | ~ spl24_260)), inference(backward_demodulation, [], [f1945, f1211])).
fof(f9900, plain, (~ spl24_72 | spl24_240 | ~ spl24_250), inference(avatar_split_clause, [], [f9899, f1876, f1826, f991])).
fof(f9899, plain, (~ (e1 = op(e2, e0)) | (spl24_240 | ~ spl24_250)), inference(forward_demodulation, [], [f1828, f1877])).
fof(f1877, plain, ((e0 = op(op(e2, op(e2, e2)), e2)) | ~ spl24_250), inference(avatar_component_clause, [], [f1876])).
fof(f9868, plain, (~ spl24_74 | ~ spl24_250 | spl24_262), inference(avatar_split_clause, [], [f9867, f1960, f1876, f999])).
fof(f9867, plain, (~ (e3 = op(e2, e0)) | (~ spl24_250 | spl24_262)), inference(forward_demodulation, [], [f1962, f1877])).
fof(f9825, plain, (~ spl24_91 | ~ spl24_97 | spl24_258), inference(avatar_contradiction_clause, [], [f9824])).
fof(f9824, plain, ($false | (~ spl24_91 | ~ spl24_97 | spl24_258)), inference(subsumption_resolution, [], [f9823, f1073])).
fof(f9823, plain, (~ (e0 = op(e1, e1)) | (~ spl24_91 | ~ spl24_97 | spl24_258)), inference(forward_demodulation, [], [f9809, f1098])).
fof(f9809, plain, (~ (e0 = op(op(e1, e0), e1)) | (~ spl24_91 | spl24_258)), inference(backward_demodulation, [], [f1934, f1073])).
fof(f1934, plain, (~ (e0 = op(op(e1, op(e1, e1)), e1)) | spl24_258), inference(avatar_component_clause, [], [f1932])).
fof(f9813, plain, (~ spl24_66 | ~ spl24_91), inference(avatar_split_clause, [], [f9799, f1071, f966])).
fof(f9799, plain, (~ (e0 = op(e2, e1)) | ~ spl24_91), inference(backward_demodulation, [], [f396, f1073])).
fof(f9795, plain, (~ spl24_18 | spl24_70 | ~ spl24_141), inference(avatar_contradiction_clause, [], [f9794])).
fof(f9794, plain, ($false | (~ spl24_18 | spl24_70 | ~ spl24_141)), inference(subsumption_resolution, [], [f9793, f983])).
fof(f983, plain, (~ (e4 = op(e2, e1)) | spl24_70), inference(avatar_component_clause, [], [f982])).
fof(f9793, plain, ((e4 = op(e2, e1)) | (~ spl24_18 | ~ spl24_141)), inference(forward_demodulation, [], [f1338, f766])).
fof(f9787, plain, (~ spl24_2 | ~ spl24_18 | spl24_220), inference(avatar_contradiction_clause, [], [f9786])).
fof(f9786, plain, ($false | (~ spl24_2 | ~ spl24_18 | spl24_220)), inference(subsumption_resolution, [], [f9785, f766])).
fof(f9785, plain, (~ (e2 = op(e4, e1)) | (~ spl24_2 | spl24_220)), inference(forward_demodulation, [], [f1725, f699])).
fof(f699, plain, ((e1 = op(e4, e4)) | ~ spl24_2), inference(avatar_component_clause, [], [f697])).
fof(f1725, plain, (~ (e2 = op(e4, op(e4, e4))) | spl24_220), inference(avatar_component_clause, [], [f1723])).
fof(f1723, plain, (spl24_220 <=> (e2 = op(e4, op(e4, e4)))), introduced(avatar_definition, [new_symbols(naming, [spl24_220])])).
fof(f9784, plain, (spl24_69 | ~ spl24_62 | ~ spl24_222), inference(avatar_split_clause, [], [f9783, f1732, f949, f978])).
fof(f9783, plain, ((e3 = op(e2, e1)) | (~ spl24_62 | ~ spl24_222)), inference(forward_demodulation, [], [f1733, f951])).
fof(f1733, plain, ((e3 = op(e2, op(e2, e2))) | ~ spl24_222), inference(avatar_component_clause, [], [f1732])).
fof(f9775, plain, (~ spl24_66 | ~ spl24_62 | spl24_253), inference(avatar_split_clause, [], [f9774, f1893, f949, f966])).
fof(f9774, plain, (~ (e0 = op(e2, e1)) | (~ spl24_62 | spl24_253)), inference(forward_demodulation, [], [f1895, f951])).
fof(f9715, plain, (~ spl24_55 | ~ spl24_60), inference(avatar_split_clause, [], [f9714, f940, f919])).
fof(f9714, plain, (~ (e4 = op(e2, e4)) | ~ spl24_60), inference(forward_demodulation, [], [f461, f942])).
fof(f9695, plain, (~ spl24_3 | ~ spl24_18), inference(avatar_split_clause, [], [f9690, f764, f701])).
fof(f9690, plain, (~ (e2 = op(e4, e4)) | ~ spl24_18), inference(backward_demodulation, [], [f478, f766])).
fof(f9659, plain, (spl24_57 | ~ spl24_83 | ~ spl24_185), inference(avatar_split_clause, [], [f9656, f1551, f1037, f928])).
fof(f9656, plain, ((e1 = op(e2, e3)) | (~ spl24_83 | ~ spl24_185)), inference(backward_demodulation, [], [f1553, f1039])).
fof(f9649, plain, (spl24_18 | ~ spl24_90 | ~ spl24_187), inference(avatar_split_clause, [], [f9645, f1561, f1066, f764])).
fof(f9645, plain, ((e2 = op(e4, e1)) | (~ spl24_90 | ~ spl24_187)), inference(backward_demodulation, [], [f1562, f1068])).
fof(f9640, plain, (spl24_84 | ~ spl24_94 | ~ spl24_232), inference(avatar_contradiction_clause, [], [f9639])).
fof(f9639, plain, ($false | (spl24_84 | ~ spl24_94 | ~ spl24_232)), inference(subsumption_resolution, [], [f9634, f1042])).
fof(f1042, plain, (~ (e3 = op(e1, e3)) | spl24_84), inference(avatar_component_clause, [], [f1041])).
fof(f9634, plain, ((e3 = op(e1, e3)) | (~ spl24_94 | ~ spl24_232)), inference(backward_demodulation, [], [f1779, f1085])).
fof(f9624, plain, (spl24_28 | ~ spl24_54 | ~ spl24_166), inference(avatar_split_clause, [], [f9623, f1457, f915, f806])).
fof(f1457, plain, (spl24_166 <=> (e2 = op(op(e2, e4), e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_166])])).
fof(f9623, plain, ((e2 = op(e3, e4)) | (~ spl24_54 | ~ spl24_166)), inference(forward_demodulation, [], [f1459, f917])).
fof(f1459, plain, ((e2 = op(op(e2, e4), e4)) | ~ spl24_166), inference(avatar_component_clause, [], [f1457])).
fof(f9618, plain, (~ spl24_6 | spl24_213 | ~ spl24_219), inference(avatar_contradiction_clause, [], [f9617])).
fof(f9617, plain, ($false | (~ spl24_6 | spl24_213 | ~ spl24_219)), inference(subsumption_resolution, [], [f9613, f716])).
fof(f9613, plain, (~ (e0 = op(e4, e3)) | (spl24_213 | ~ spl24_219)), inference(backward_demodulation, [], [f1695, f1720])).
fof(f1720, plain, ((e3 = op(op(e4, op(e4, e4)), e4)) | ~ spl24_219), inference(avatar_component_clause, [], [f1719])).
fof(f9512, plain, (~ spl24_7 | ~ spl24_57), inference(avatar_split_clause, [], [f9505, f928, f718])).
fof(f9505, plain, (~ (e1 = op(e4, e3)) | ~ spl24_57), inference(backward_demodulation, [], [f420, f930])).
fof(f9504, plain, (~ spl24_51 | ~ spl24_65 | spl24_253), inference(avatar_split_clause, [], [f9501, f1893, f961, f903])).
fof(f9501, plain, (~ (e0 = op(e2, e4)) | (~ spl24_65 | spl24_253)), inference(backward_demodulation, [], [f1895, f963])).
fof(f9502, plain, (~ spl24_60 | ~ spl24_65), inference(avatar_split_clause, [], [f9497, f961, f940])).
fof(f9497, plain, (~ (e4 = op(e2, e3)) | ~ spl24_65), inference(backward_demodulation, [], [f459, f963])).
fof(f9494, plain, (~ spl24_51 | ~ spl24_66), inference(avatar_split_clause, [], [f9486, f966, f903])).
fof(f9486, plain, (~ (e0 = op(e2, e4)) | ~ spl24_66), inference(backward_demodulation, [], [f458, f968])).
fof(f9493, plain, (~ spl24_16 | ~ spl24_66), inference(avatar_split_clause, [], [f9483, f966, f756])).
fof(f9483, plain, (~ (e0 = op(e4, e1)) | ~ spl24_66), inference(backward_demodulation, [], [f400, f968])).
fof(f9461, plain, (~ spl24_76 | ~ spl24_45 | ~ spl24_89 | ~ spl24_93 | spl24_227), inference(avatar_split_clause, [], [f9460, f1756, f1079, f1062, f877, f1008])).
fof(f9460, plain, (~ (e0 = op(e1, e4)) | (~ spl24_45 | ~ spl24_89 | ~ spl24_93 | spl24_227)), inference(forward_demodulation, [], [f9451, f879])).
fof(f9451, plain, (~ (e0 = op(e1, op(e3, e1))) | (~ spl24_89 | ~ spl24_93 | spl24_227)), inference(backward_demodulation, [], [f9426, f1064])).
fof(f9426, plain, (~ (e0 = op(e1, op(op(e1, e2), e1))) | (~ spl24_93 | spl24_227)), inference(forward_demodulation, [], [f1758, f1081])).
fof(f9434, plain, (~ spl24_7 | ~ spl24_45 | spl24_157), inference(avatar_split_clause, [], [f9433, f1414, f877, f718])).
fof(f9433, plain, (~ (e1 = op(e4, e3)) | (~ spl24_45 | spl24_157)), inference(forward_demodulation, [], [f1416, f879])).
fof(f1416, plain, (~ (e1 = op(op(e3, e1), e3)) | spl24_157), inference(avatar_component_clause, [], [f1414])).
fof(f9432, plain, (spl24_19 | ~ spl24_45 | ~ spl24_158), inference(avatar_split_clause, [], [f9431, f1419, f877, f768])).
fof(f9431, plain, ((e3 = op(e4, e1)) | (~ spl24_45 | ~ spl24_158)), inference(forward_demodulation, [], [f1421, f879])).
fof(f1421, plain, ((e3 = op(op(e3, e1), e1)) | ~ spl24_158), inference(avatar_component_clause, [], [f1419])).
fof(f9425, plain, (~ spl24_90 | ~ spl24_93 | spl24_229), inference(avatar_split_clause, [], [f9424, f1764, f1079, f1066])).
fof(f9424, plain, (~ (e4 = op(e1, e2)) | (~ spl24_93 | spl24_229)), inference(forward_demodulation, [], [f1766, f1081])).
fof(f9421, plain, (~ spl24_36 | ~ spl24_222 | spl24_250), inference(avatar_contradiction_clause, [], [f9420])).
fof(f9420, plain, ($false | (~ spl24_36 | ~ spl24_222 | spl24_250)), inference(subsumption_resolution, [], [f9419, f842])).
fof(f9419, plain, (~ (e0 = op(e3, e2)) | (~ spl24_222 | spl24_250)), inference(forward_demodulation, [], [f1878, f1733])).
fof(f9403, plain, (~ spl24_82 | ~ spl24_97), inference(avatar_split_clause, [], [f9402, f1096, f1033])).
fof(f9402, plain, (~ (e1 = op(e1, e3)) | ~ spl24_97), inference(forward_demodulation, [], [f444, f1098])).
fof(f9399, plain, (~ spl24_85 | ~ spl24_27 | spl24_148), inference(avatar_split_clause, [], [f9100, f1369, f802, f1045])).
fof(f9100, plain, (~ (e4 = op(e1, e3)) | (~ spl24_27 | spl24_148)), inference(backward_demodulation, [], [f1371, f804])).
fof(f9388, plain, (~ spl24_20 | ~ spl24_45), inference(avatar_split_clause, [], [f9387, f877, f772])).
fof(f9387, plain, (~ (e4 = op(e4, e1)) | ~ spl24_45), inference(forward_demodulation, [], [f401, f879])).
fof(f9385, plain, (~ spl24_18 | ~ spl24_27 | ~ spl24_215 | spl24_255), inference(avatar_split_clause, [], [f9294, f1909, f1701, f802, f764])).
fof(f9294, plain, (~ (e2 = op(e4, e1)) | (~ spl24_27 | ~ spl24_215 | spl24_255)), inference(forward_demodulation, [], [f9288, f804])).
fof(f9288, plain, (~ (e2 = op(e4, op(e3, e4))) | (~ spl24_215 | spl24_255)), inference(backward_demodulation, [], [f1911, f1702])).
fof(f9381, plain, (~ spl24_12 | ~ spl24_3 | spl24_230), inference(avatar_split_clause, [], [f9378, f1769, f701, f739])).
fof(f9378, plain, (~ (e1 = op(e4, e2)) | (~ spl24_3 | spl24_230)), inference(backward_demodulation, [], [f1771, f703])).
fof(f9319, plain, (~ spl24_89 | ~ spl24_93 | spl24_232), inference(avatar_split_clause, [], [f9313, f1778, f1079, f1062])).
fof(f9313, plain, (~ (e3 = op(e1, e2)) | (~ spl24_93 | spl24_232)), inference(backward_demodulation, [], [f1780, f1081])).
fof(f9318, plain, (spl24_90 | ~ spl24_93 | ~ spl24_229), inference(avatar_split_clause, [], [f9312, f1764, f1079, f1066])).
fof(f9312, plain, ((e4 = op(e1, e2)) | (~ spl24_93 | ~ spl24_229)), inference(backward_demodulation, [], [f1765, f1081])).
fof(f9317, plain, (~ spl24_78 | ~ spl24_93), inference(avatar_split_clause, [], [f9309, f1079, f1016])).
fof(f9309, plain, (~ (e2 = op(e1, e4)) | ~ spl24_93), inference(backward_demodulation, [], [f448, f1081])).
fof(f9316, plain, (~ spl24_43 | ~ spl24_93), inference(avatar_split_clause, [], [f9307, f1079, f869])).
fof(f9307, plain, (~ (e2 = op(e3, e1)) | ~ spl24_93), inference(backward_demodulation, [], [f397, f1081])).
fof(f9291, plain, (~ spl24_27 | ~ spl24_215 | spl24_226), inference(avatar_contradiction_clause, [], [f9290])).
fof(f9290, plain, ($false | (~ spl24_27 | ~ spl24_215 | spl24_226)), inference(subsumption_resolution, [], [f9285, f804])).
fof(f9285, plain, (~ (e1 = op(e3, e4)) | (~ spl24_215 | spl24_226)), inference(backward_demodulation, [], [f1753, f1702])).
fof(f9244, plain, (spl24_43 | ~ spl24_69 | ~ spl24_174), inference(avatar_split_clause, [], [f9243, f1498, f978, f869])).
fof(f9243, plain, ((e2 = op(e3, e1)) | (~ spl24_69 | ~ spl24_174)), inference(forward_demodulation, [], [f1500, f980])).
fof(f9237, plain, (~ spl24_59 | ~ spl24_69), inference(avatar_split_clause, [], [f9236, f978, f936])).
fof(f9236, plain, (~ (e3 = op(e2, e3)) | ~ spl24_69), inference(forward_demodulation, [], [f457, f980])).
fof(f9232, plain, (~ spl24_44 | ~ spl24_69), inference(avatar_split_clause, [], [f9231, f978, f873])).
fof(f9231, plain, (~ (e3 = op(e3, e1)) | ~ spl24_69), inference(forward_demodulation, [], [f399, f980])).
fof(f9196, plain, (spl24_66 | ~ spl24_78 | ~ spl24_95 | ~ spl24_258), inference(avatar_contradiction_clause, [], [f9195])).
fof(f9195, plain, ($false | (spl24_66 | ~ spl24_78 | ~ spl24_95 | ~ spl24_258)), inference(subsumption_resolution, [], [f9194, f967])).
fof(f967, plain, (~ (e0 = op(e2, e1)) | spl24_66), inference(avatar_component_clause, [], [f966])).
fof(f9194, plain, ((e0 = op(e2, e1)) | (~ spl24_78 | ~ spl24_95 | ~ spl24_258)), inference(forward_demodulation, [], [f9187, f1018])).
fof(f9187, plain, ((e0 = op(op(e1, e4), e1)) | (~ spl24_95 | ~ spl24_258)), inference(backward_demodulation, [], [f1933, f1089])).
fof(f9190, plain, (~ spl24_90 | ~ spl24_95), inference(avatar_split_clause, [], [f9184, f1087, f1066])).
fof(f9184, plain, (~ (e4 = op(e1, e2)) | ~ spl24_95), inference(backward_demodulation, [], [f446, f1089])).
fof(f9179, plain, (spl24_69 | ~ spl24_43 | ~ spl24_158), inference(avatar_split_clause, [], [f9178, f1419, f869, f978])).
fof(f9178, plain, ((e3 = op(e2, e1)) | (~ spl24_43 | ~ spl24_158)), inference(forward_demodulation, [], [f1421, f871])).
fof(f9177, plain, (~ spl24_51 | spl24_103 | ~ spl24_166), inference(avatar_contradiction_clause, [], [f9176])).
fof(f9176, plain, ($false | (~ spl24_51 | spl24_103 | ~ spl24_166)), inference(subsumption_resolution, [], [f9175, f1122])).
fof(f1122, plain, (~ (e2 = op(e0, e4)) | spl24_103), inference(avatar_component_clause, [], [f1121])).
fof(f9175, plain, ((e2 = op(e0, e4)) | (~ spl24_51 | ~ spl24_166)), inference(forward_demodulation, [], [f1459, f905])).
fof(f9157, plain, (~ spl24_51 | spl24_216 | ~ spl24_221), inference(avatar_contradiction_clause, [], [f9156])).
fof(f9156, plain, ($false | (~ spl24_51 | spl24_216 | ~ spl24_221)), inference(subsumption_resolution, [], [f9155, f905])).
fof(f9155, plain, (~ (e0 = op(e2, e4)) | (spl24_216 | ~ spl24_221)), inference(backward_demodulation, [], [f1708, f1729])).
fof(f1729, plain, ((e4 = op(op(e2, op(e2, e2)), e2)) | ~ spl24_221), inference(avatar_component_clause, [], [f1728])).
fof(f9134, plain, (~ spl24_1 | ~ spl24_4), inference(avatar_contradiction_clause, [], [f9133])).
fof(f9133, plain, ($false | (~ spl24_1 | ~ spl24_4)), inference(subsumption_resolution, [], [f9132, f484])).
fof(f9132, plain, ((e0 = e3) | (~ spl24_1 | ~ spl24_4)), inference(backward_demodulation, [], [f707, f695])).
fof(f9112, plain, (~ spl24_11 | ~ spl24_16), inference(avatar_split_clause, [], [f9105, f756, f735])).
fof(f9105, plain, (~ (e0 = op(e4, e2)) | ~ spl24_16), inference(backward_demodulation, [], [f476, f758])).
fof(f9095, plain, (~ spl24_26 | ~ spl24_35 | spl24_249), inference(avatar_split_clause, [], [f9093, f1871, f835, f798])).
fof(f9093, plain, (~ (e0 = op(e3, e4)) | (~ spl24_35 | spl24_249)), inference(backward_demodulation, [], [f1873, f837])).
fof(f9094, plain, (spl24_27 | ~ spl24_35 | ~ spl24_233), inference(avatar_split_clause, [], [f9092, f1783, f835, f802])).
fof(f9092, plain, ((e1 = op(e3, e4)) | (~ spl24_35 | ~ spl24_233)), inference(backward_demodulation, [], [f1784, f837])).
fof(f9085, plain, (~ spl24_11 | ~ spl24_51 | spl24_138), inference(avatar_split_clause, [], [f9080, f1322, f903, f735])).
fof(f9080, plain, (~ (e0 = op(e4, e2)) | (~ spl24_51 | spl24_138)), inference(backward_demodulation, [], [f1324, f905])).
fof(f9048, plain, (~ spl24_12 | ~ spl24_70 | spl24_173), inference(avatar_split_clause, [], [f9045, f1493, f982, f739])).
fof(f9045, plain, (~ (e1 = op(e4, e2)) | (~ spl24_70 | spl24_173)), inference(backward_demodulation, [], [f1495, f984])).
fof(f9037, plain, (spl24_12 | ~ spl24_90 | ~ spl24_188), inference(avatar_split_clause, [], [f9032, f1566, f1066, f739])).
fof(f9032, plain, ((e1 = op(e4, e2)) | (~ spl24_90 | ~ spl24_188)), inference(backward_demodulation, [], [f1568, f1068])).
fof(f9035, plain, (~ spl24_40 | ~ spl24_90), inference(avatar_split_clause, [], [f9028, f1066, f856])).
fof(f9028, plain, (~ (e4 = op(e3, e2)) | ~ spl24_90), inference(backward_demodulation, [], [f407, f1068])).
fof(f9019, plain, (~ spl24_54 | ~ spl24_8 | spl24_132), inference(avatar_split_clause, [], [f9018, f1293, f722, f915])).
fof(f9018, plain, (~ (e3 = op(e2, e4)) | (~ spl24_8 | spl24_132)), inference(forward_demodulation, [], [f1295, f724])).
fof(f8966, plain, (~ spl24_68 | ~ spl24_73), inference(avatar_split_clause, [], [f8965, f995, f974])).
fof(f8965, plain, (~ (e2 = op(e2, e1)) | ~ spl24_73), inference(forward_demodulation, [], [f452, f997])).
fof(f997, plain, ((e2 = op(e2, e0)) | ~ spl24_73), inference(avatar_component_clause, [], [f995])).
fof(f8929, plain, (~ spl24_26 | ~ spl24_60 | spl24_210 | ~ spl24_224), inference(avatar_split_clause, [], [f8923, f1741, f1680, f940, f798])).
fof(f8923, plain, (~ (e0 = op(e3, e4)) | (~ spl24_60 | spl24_210 | ~ spl24_224)), inference(backward_demodulation, [], [f8885, f942])).
fof(f8885, plain, (~ (e0 = op(e3, op(e2, e3))) | (spl24_210 | ~ spl24_224)), inference(forward_demodulation, [], [f1682, f1742])).
fof(f8906, plain, (spl24_60 | ~ spl24_8 | ~ spl24_133), inference(avatar_split_clause, [], [f8905, f1298, f722, f940])).
fof(f8905, plain, ((e4 = op(e2, e3)) | (~ spl24_8 | ~ spl24_133)), inference(forward_demodulation, [], [f1300, f724])).
fof(f8902, plain, (~ spl24_19 | ~ spl24_43 | ~ spl24_141), inference(avatar_contradiction_clause, [], [f8901])).
fof(f8901, plain, ($false | (~ spl24_19 | ~ spl24_43 | ~ spl24_141)), inference(subsumption_resolution, [], [f8900, f490])).
fof(f8900, plain, ((e2 = e4) | (~ spl24_19 | ~ spl24_43 | ~ spl24_141)), inference(forward_demodulation, [], [f8899, f871])).
fof(f8899, plain, ((e4 = op(e3, e1)) | (~ spl24_19 | ~ spl24_141)), inference(forward_demodulation, [], [f1338, f770])).
fof(f8891, plain, (~ spl24_57 | ~ spl24_43 | spl24_157), inference(avatar_split_clause, [], [f8890, f1414, f869, f928])).
fof(f8890, plain, (~ (e1 = op(e2, e3)) | (~ spl24_43 | spl24_157)), inference(forward_demodulation, [], [f1416, f871])).
fof(f8865, plain, (~ spl24_5 | ~ spl24_25), inference(avatar_split_clause, [], [f8864, f793, f709])).
fof(f8864, plain, (~ (e4 = op(e4, e4)) | ~ spl24_25), inference(forward_demodulation, [], [f475, f795])).
fof(f8859, plain, (~ spl24_4 | ~ spl24_19), inference(avatar_split_clause, [], [f8858, f768, f705])).
fof(f8858, plain, (~ (e3 = op(e4, e4)) | ~ spl24_19), inference(forward_demodulation, [], [f478, f770])).
fof(f8854, plain, (~ spl24_12 | ~ spl24_89 | ~ spl24_137), inference(avatar_contradiction_clause, [], [f8853])).
fof(f8853, plain, ($false | (~ spl24_12 | ~ spl24_89 | ~ spl24_137)), inference(subsumption_resolution, [], [f8852, f491])).
fof(f491, plain, ~ (e3 = e4), inference(cnf_transformation, [], [f5])).
fof(f8852, plain, ((e3 = e4) | (~ spl24_12 | ~ spl24_89 | ~ spl24_137)), inference(forward_demodulation, [], [f8849, f1064])).
fof(f8849, plain, ((e4 = op(e1, e2)) | (~ spl24_12 | ~ spl24_137)), inference(backward_demodulation, [], [f1319, f741])).
fof(f8823, plain, (~ spl24_27 | ~ spl24_35 | spl24_233), inference(avatar_split_clause, [], [f8818, f1783, f835, f802])).
fof(f8818, plain, (~ (e1 = op(e3, e4)) | (~ spl24_35 | spl24_233)), inference(backward_demodulation, [], [f1785, f837])).
fof(f8810, plain, (~ spl24_26 | ~ spl24_36), inference(avatar_split_clause, [], [f8802, f840, f798])).
fof(f8802, plain, (~ (e0 = op(e3, e4)) | ~ spl24_36), inference(backward_demodulation, [], [f470, f842])).
fof(f8809, plain, (~ spl24_11 | ~ spl24_36), inference(avatar_split_clause, [], [f8800, f840, f735])).
fof(f8800, plain, (~ (e0 = op(e4, e2)) | ~ spl24_36), inference(backward_demodulation, [], [f411, f842])).
fof(f8798, plain, (~ spl24_57 | ~ spl24_89 | spl24_168), inference(avatar_contradiction_clause, [], [f8797])).
fof(f8797, plain, ($false | (~ spl24_57 | ~ spl24_89 | spl24_168)), inference(subsumption_resolution, [], [f8794, f1064])).
fof(f8794, plain, (~ (e3 = op(e1, e2)) | (~ spl24_57 | spl24_168)), inference(backward_demodulation, [], [f1469, f930])).
fof(f8767, plain, (~ spl24_66 | ~ spl24_78 | ~ spl24_95 | spl24_258), inference(avatar_contradiction_clause, [], [f8766])).
fof(f8766, plain, ($false | (~ spl24_66 | ~ spl24_78 | ~ spl24_95 | spl24_258)), inference(subsumption_resolution, [], [f8765, f968])).
fof(f8765, plain, (~ (e0 = op(e2, e1)) | (~ spl24_78 | ~ spl24_95 | spl24_258)), inference(forward_demodulation, [], [f8764, f1018])).
fof(f8764, plain, (~ (e0 = op(op(e1, e4), e1)) | (~ spl24_95 | spl24_258)), inference(forward_demodulation, [], [f1934, f1089])).
fof(f8725, plain, (~ spl24_1 | ~ spl24_2), inference(avatar_contradiction_clause, [], [f8724])).
fof(f8724, plain, ($false | (~ spl24_1 | ~ spl24_2)), inference(subsumption_resolution, [], [f8723, f482])).
fof(f8723, plain, ((e0 = e1) | (~ spl24_1 | ~ spl24_2)), inference(forward_demodulation, [], [f699, f695])).
fof(f8714, plain, (~ spl24_11 | ~ spl24_113 | ~ spl24_137), inference(avatar_contradiction_clause, [], [f8713])).
fof(f8713, plain, ($false | (~ spl24_11 | ~ spl24_113 | ~ spl24_137)), inference(subsumption_resolution, [], [f8712, f490])).
fof(f8712, plain, ((e2 = e4) | (~ spl24_11 | ~ spl24_113 | ~ spl24_137)), inference(forward_demodulation, [], [f8711, f1165])).
fof(f8711, plain, ((e4 = op(e0, e2)) | (~ spl24_11 | ~ spl24_137)), inference(backward_demodulation, [], [f1319, f737])).
fof(f8686, plain, (~ spl24_32 | ~ spl24_43 | spl24_224), inference(avatar_contradiction_clause, [], [f8685])).
fof(f8685, plain, ($false | (~ spl24_32 | ~ spl24_43 | spl24_224)), inference(subsumption_resolution, [], [f8675, f871])).
fof(f8675, plain, (~ (e2 = op(e3, e1)) | (~ spl24_32 | spl24_224)), inference(backward_demodulation, [], [f1743, f825])).
fof(f8680, plain, (~ spl24_7 | ~ spl24_32), inference(avatar_split_clause, [], [f8667, f823, f718])).
fof(f8667, plain, (~ (e1 = op(e4, e3)) | ~ spl24_32), inference(backward_demodulation, [], [f421, f825])).
fof(f8657, plain, (~ spl24_18 | ~ spl24_43), inference(avatar_split_clause, [], [f8649, f869, f764])).
fof(f8649, plain, (~ (e2 = op(e4, e1)) | ~ spl24_43), inference(backward_demodulation, [], [f401, f871])).
fof(f8638, plain, (~ spl24_45 | ~ spl24_95), inference(avatar_split_clause, [], [f8634, f1087, f877])).
fof(f8634, plain, (~ (e4 = op(e3, e1)) | ~ spl24_95), inference(backward_demodulation, [], [f397, f1089])).
fof(f8624, plain, (~ spl24_27 | ~ spl24_223 | spl24_238), inference(avatar_split_clause, [], [f8622, f1816, f1737, f802])).
fof(f8622, plain, (~ (e1 = op(e3, e4)) | (~ spl24_223 | spl24_238)), inference(backward_demodulation, [], [f1818, f1738])).
fof(f8601, plain, (~ spl24_39 | ~ spl24_49), inference(avatar_split_clause, [], [f8600, f894, f852])).
fof(f8600, plain, (~ (e3 = op(e3, e2)) | ~ spl24_49), inference(forward_demodulation, [], [f463, f896])).
fof(f896, plain, ((e3 = op(e3, e0)) | ~ spl24_49), inference(avatar_component_clause, [], [f894])).
fof(f8599, plain, (~ spl24_38 | ~ spl24_113), inference(avatar_split_clause, [], [f8598, f1163, f848])).
fof(f8598, plain, (~ (e2 = op(e3, e2)) | ~ spl24_113), inference(forward_demodulation, [], [f404, f1165])).
fof(f8589, plain, (~ spl24_15 | ~ spl24_25), inference(avatar_split_clause, [], [f8588, f793, f751])).
fof(f8588, plain, (~ (e4 = op(e4, e2)) | ~ spl24_25), inference(forward_demodulation, [], [f473, f795])).
fof(f8537, plain, (~ spl24_14 | ~ spl24_36 | ~ spl24_137), inference(avatar_contradiction_clause, [], [f8536])).
fof(f8536, plain, ($false | (~ spl24_14 | ~ spl24_36 | ~ spl24_137)), inference(subsumption_resolution, [], [f8535, f485])).
fof(f8535, plain, ((e0 = e4) | (~ spl24_14 | ~ spl24_36 | ~ spl24_137)), inference(forward_demodulation, [], [f8534, f842])).
fof(f8534, plain, ((e4 = op(e3, e2)) | (~ spl24_14 | ~ spl24_137)), inference(forward_demodulation, [], [f1319, f749])).
fof(f8529, plain, (~ spl24_18 | ~ spl24_66 | ~ spl24_141), inference(avatar_contradiction_clause, [], [f8528])).
fof(f8528, plain, ($false | (~ spl24_18 | ~ spl24_66 | ~ spl24_141)), inference(subsumption_resolution, [], [f8527, f485])).
fof(f8527, plain, ((e0 = e4) | (~ spl24_18 | ~ spl24_66 | ~ spl24_141)), inference(forward_demodulation, [], [f8526, f968])).
fof(f8526, plain, ((e4 = op(e2, e1)) | (~ spl24_18 | ~ spl24_141)), inference(backward_demodulation, [], [f1338, f766])).
fof(f8482, plain, (~ spl24_52 | ~ spl24_78 | spl24_182), inference(avatar_split_clause, [], [f8479, f1536, f1016, f907])).
fof(f8479, plain, (~ (e1 = op(e2, e4)) | (~ spl24_78 | spl24_182)), inference(backward_demodulation, [], [f1537, f1018])).
fof(f1537, plain, (~ (e1 = op(op(e1, e4), e4)) | spl24_182), inference(avatar_component_clause, [], [f1536])).
fof(f8380, plain, (~ spl24_28 | ~ spl24_36 | ~ spl24_134 | spl24_210 | ~ spl24_212), inference(avatar_contradiction_clause, [], [f8379])).
fof(f8379, plain, ($false | (~ spl24_28 | ~ spl24_36 | ~ spl24_134 | spl24_210 | ~ spl24_212)), inference(subsumption_resolution, [], [f8371, f842])).
fof(f8371, plain, (~ (e0 = op(e3, e2)) | (~ spl24_28 | ~ spl24_134 | spl24_210 | ~ spl24_212)), inference(backward_demodulation, [], [f8322, f808])).
fof(f8322, plain, (~ (e0 = op(e3, op(e3, e4))) | (~ spl24_134 | spl24_210 | ~ spl24_212)), inference(forward_demodulation, [], [f8317, f1304])).
fof(f8317, plain, (~ (e0 = op(e3, op(e4, e3))) | (spl24_210 | ~ spl24_212)), inference(backward_demodulation, [], [f1682, f1689])).
fof(f8343, plain, (~ spl24_25 | ~ spl24_105 | spl24_146), inference(avatar_contradiction_clause, [], [f8342])).
fof(f8342, plain, ($false | (~ spl24_25 | ~ spl24_105 | spl24_146)), inference(subsumption_resolution, [], [f8341, f1131])).
fof(f8341, plain, (~ (e4 = op(e0, e4)) | (~ spl24_25 | spl24_146)), inference(forward_demodulation, [], [f1362, f795])).
fof(f8237, plain, (~ spl24_1 | ~ spl24_3), inference(avatar_contradiction_clause, [], [f8236])).
fof(f8236, plain, ($false | (~ spl24_1 | ~ spl24_3)), inference(subsumption_resolution, [], [f8235, f483])).
fof(f8235, plain, ((e0 = e2) | (~ spl24_1 | ~ spl24_3)), inference(backward_demodulation, [], [f703, f695])).
fof(f8187, plain, (~ spl24_52 | spl24_78 | ~ spl24_166), inference(avatar_contradiction_clause, [], [f8186])).
fof(f8186, plain, ($false | (~ spl24_52 | spl24_78 | ~ spl24_166)), inference(subsumption_resolution, [], [f8179, f1017])).
fof(f1017, plain, (~ (e2 = op(e1, e4)) | spl24_78), inference(avatar_component_clause, [], [f1016])).
fof(f8179, plain, ((e2 = op(e1, e4)) | (~ spl24_52 | ~ spl24_166)), inference(backward_demodulation, [], [f1459, f909])).
fof(f8185, plain, (~ spl24_52 | ~ spl24_90 | spl24_165), inference(avatar_contradiction_clause, [], [f8184])).
fof(f8184, plain, ($false | (~ spl24_52 | ~ spl24_90 | spl24_165)), inference(subsumption_resolution, [], [f8178, f1068])).
fof(f8178, plain, (~ (e4 = op(e1, e2)) | (~ spl24_52 | spl24_165)), inference(backward_demodulation, [], [f1454, f909])).
fof(f8183, plain, (spl24_12 | ~ spl24_52 | ~ spl24_138), inference(avatar_split_clause, [], [f8177, f1322, f907, f739])).
fof(f8177, plain, ((e1 = op(e4, e2)) | (~ spl24_52 | ~ spl24_138)), inference(backward_demodulation, [], [f1323, f909])).
fof(f8182, plain, (~ spl24_27 | ~ spl24_52), inference(avatar_split_clause, [], [f8172, f907, f802])).
fof(f8172, plain, (~ (e1 = op(e3, e4)) | ~ spl24_52), inference(backward_demodulation, [], [f429, f909])).
fof(f8162, plain, (~ spl24_53 | ~ spl24_73), inference(avatar_split_clause, [], [f8157, f995, f911])).
fof(f8157, plain, (~ (e2 = op(e2, e4)) | ~ spl24_73), inference(backward_demodulation, [], [f455, f997])).
fof(f8152, plain, (~ spl24_79 | ~ spl24_81 | ~ spl24_90 | ~ spl24_93 | ~ spl24_142 | spl24_227), inference(avatar_contradiction_clause, [], [f8151])).
fof(f8151, plain, ($false | (~ spl24_79 | ~ spl24_81 | ~ spl24_90 | ~ spl24_93 | ~ spl24_142 | spl24_227)), inference(subsumption_resolution, [], [f8143, f1031])).
fof(f1031, plain, ((e0 = op(e1, e3)) | ~ spl24_81), inference(avatar_component_clause, [], [f1029])).
fof(f8143, plain, (~ (e0 = op(e1, e3)) | (~ spl24_79 | ~ spl24_90 | ~ spl24_93 | ~ spl24_142 | spl24_227)), inference(backward_demodulation, [], [f8130, f1022])).
fof(f8130, plain, (~ (e0 = op(e1, op(e1, e4))) | (~ spl24_90 | ~ spl24_93 | ~ spl24_142 | spl24_227)), inference(forward_demodulation, [], [f8122, f1342])).
fof(f8122, plain, (~ (e0 = op(e1, op(e4, e1))) | (~ spl24_90 | ~ spl24_93 | spl24_227)), inference(backward_demodulation, [], [f7376, f1068])).
fof(f7376, plain, (~ (e0 = op(e1, op(op(e1, e2), e1))) | (~ spl24_93 | spl24_227)), inference(backward_demodulation, [], [f1758, f1081])).
fof(f8147, plain, (~ spl24_29 | ~ spl24_79), inference(avatar_split_clause, [], [f8135, f1020, f810])).
fof(f8135, plain, (~ (e3 = op(e3, e4)) | ~ spl24_79), inference(backward_demodulation, [], [f427, f1022])).
fof(f8116, plain, (~ spl24_77 | ~ spl24_97), inference(avatar_split_clause, [], [f8110, f1096, f1012])).
fof(f8110, plain, (~ (e1 = op(e1, e4)) | ~ spl24_97), inference(backward_demodulation, [], [f445, f1098])).
fof(f8105, plain, (~ spl24_25 | ~ spl24_105 | spl24_196), inference(avatar_split_clause, [], [f8096, f1606, f1129, f793])).
fof(f8096, plain, (~ (e4 = op(e4, e0)) | (~ spl24_105 | spl24_196)), inference(backward_demodulation, [], [f1608, f1131])).
fof(f8103, plain, (~ spl24_5 | ~ spl24_105), inference(avatar_split_clause, [], [f8094, f1129, f709])).
fof(f8094, plain, (~ (e4 = op(e4, e4)) | ~ spl24_105), inference(backward_demodulation, [], [f425, f1131])).
fof(f8089, plain, (~ spl24_74 | ~ spl24_109 | ~ spl24_163), inference(avatar_split_clause, [], [f8081, f1443, f1146, f999])).
fof(f8081, plain, (~ (e3 = op(e2, e0)) | (~ spl24_109 | ~ spl24_163)), inference(backward_demodulation, [], [f7444, f1148])).
fof(f7444, plain, (~ (op(e0, e3) = op(e2, e0)) | ~ spl24_163), inference(backward_demodulation, [], [f389, f1444])).
fof(f8086, plain, (~ spl24_9 | ~ spl24_109), inference(avatar_split_clause, [], [f8075, f1146, f726])).
fof(f8075, plain, (~ (e3 = op(e4, e3)) | ~ spl24_109), inference(backward_demodulation, [], [f415, f1148])).
fof(f8072, plain, (~ spl24_73 | ~ spl24_113 | spl24_179), inference(avatar_split_clause, [], [f8067, f1522, f1163, f995])).
fof(f8067, plain, (~ (e2 = op(e2, e0)) | (~ spl24_113 | spl24_179)), inference(backward_demodulation, [], [f1524, f1165])).
fof(f8070, plain, (~ spl24_13 | ~ spl24_113), inference(avatar_split_clause, [], [f8062, f1163, f743])).
fof(f8062, plain, (~ (e2 = op(e4, e2)) | ~ spl24_113), inference(backward_demodulation, [], [f405, f1165])).
fof(f8059, plain, (~ spl24_97 | ~ spl24_117 | spl24_205), inference(avatar_split_clause, [], [f8055, f1651, f1180, f1096])).
fof(f8055, plain, (~ (e1 = op(e1, e0)) | (~ spl24_117 | spl24_205)), inference(backward_demodulation, [], [f1653, f1182])).
fof(f8058, plain, (~ spl24_97 | ~ spl24_117 | spl24_194), inference(avatar_split_clause, [], [f8054, f1597, f1180, f1096])).
fof(f8054, plain, (~ (e1 = op(e1, e0)) | (~ spl24_117 | spl24_194)), inference(backward_demodulation, [], [f1599, f1182])).
fof(f8057, plain, (~ spl24_17 | ~ spl24_117), inference(avatar_split_clause, [], [f8047, f1180, f760])).
fof(f8047, plain, (~ (e1 = op(e4, e1)) | ~ spl24_117), inference(backward_demodulation, [], [f395, f1182])).
fof(f8036, plain, (~ spl24_101 | ~ spl24_121), inference(avatar_split_clause, [], [f8024, f1197, f1113])).
fof(f8024, plain, (~ (e0 = op(e0, e4)) | ~ spl24_121), inference(backward_demodulation, [], [f435, f1199])).
fof(f7985, plain, (~ spl24_119 | ~ spl24_81 | spl24_184), inference(avatar_split_clause, [], [f7984, f1546, f1029, f1188])).
fof(f7984, plain, (~ (e3 = op(e0, e1)) | (~ spl24_81 | spl24_184)), inference(forward_demodulation, [], [f1548, f1031])).
fof(f7979, plain, (~ spl24_12 | ~ spl24_218 | spl24_234), inference(avatar_split_clause, [], [f7978, f1788, f1714, f739])).
fof(f7978, plain, (~ (e1 = op(e4, e2)) | (~ spl24_218 | spl24_234)), inference(forward_demodulation, [], [f1790, f1715])).
fof(f1715, plain, ((e4 = op(e2, op(e2, e2))) | ~ spl24_218), inference(avatar_component_clause, [], [f1714])).
fof(f7876, plain, (~ spl24_48 | ~ spl24_115 | ~ spl24_119 | ~ spl24_122 | spl24_269), inference(avatar_contradiction_clause, [], [f7875])).
fof(f7875, plain, ($false | (~ spl24_48 | ~ spl24_115 | ~ spl24_119 | ~ spl24_122 | spl24_269)), inference(subsumption_resolution, [], [f7874, f1173])).
fof(f7874, plain, (~ (e4 = op(e0, e2)) | (~ spl24_48 | ~ spl24_119 | ~ spl24_122 | spl24_269)), inference(forward_demodulation, [], [f7873, f892])).
fof(f7873, plain, (~ (e4 = op(e0, op(e3, e0))) | (~ spl24_119 | ~ spl24_122 | spl24_269)), inference(forward_demodulation, [], [f7872, f1190])).
fof(f7872, plain, (~ (e4 = op(e0, op(op(e0, e1), e0))) | (~ spl24_122 | spl24_269)), inference(forward_demodulation, [], [f2022, f1203])).
fof(f2022, plain, (~ (e4 = op(e0, op(op(e0, op(e0, e0)), e0))) | spl24_269), inference(avatar_component_clause, [], [f2020])).
fof(f2020, plain, (spl24_269 <=> (e4 = op(e0, op(op(e0, op(e0, e0)), e0)))), introduced(avatar_definition, [new_symbols(naming, [spl24_269])])).
fof(f7831, plain, (spl24_112 | ~ spl24_66 | ~ spl24_173), inference(avatar_split_clause, [], [f7334, f1493, f966, f1159])).
fof(f7334, plain, ((e1 = op(e0, e2)) | (~ spl24_66 | ~ spl24_173)), inference(forward_demodulation, [], [f1494, f968])).
fof(f1494, plain, ((e1 = op(op(e2, e1), e2)) | ~ spl24_173), inference(avatar_component_clause, [], [f1493])).
fof(f7801, plain, (spl24_48 | ~ spl24_74 | ~ spl24_178), inference(avatar_split_clause, [], [f7797, f1517, f999, f890])).
fof(f7797, plain, ((e2 = op(e3, e0)) | (~ spl24_74 | ~ spl24_178)), inference(backward_demodulation, [], [f1519, f1001])).
fof(f7798, plain, (~ spl24_64 | ~ spl24_74), inference(avatar_split_clause, [], [f7793, f999, f957])).
fof(f7793, plain, (~ (e3 = op(e2, e2)) | ~ spl24_74), inference(backward_demodulation, [], [f453, f1001])).
fof(f7773, plain, (~ spl24_90 | ~ spl24_100), inference(avatar_split_clause, [], [f7769, f1108, f1066])).
fof(f7769, plain, (~ (e4 = op(e1, e2)) | ~ spl24_100), inference(backward_demodulation, [], [f443, f1110])).
fof(f7772, plain, (~ spl24_75 | ~ spl24_100), inference(avatar_split_clause, [], [f7768, f1108, f1003])).
fof(f7768, plain, (~ (e4 = op(e2, e0)) | ~ spl24_100), inference(backward_demodulation, [], [f386, f1110])).
fof(f7755, plain, (~ spl24_112 | ~ spl24_115), inference(avatar_contradiction_clause, [], [f7754])).
fof(f7754, plain, ($false | (~ spl24_112 | ~ spl24_115)), inference(subsumption_resolution, [], [f7753, f488])).
fof(f7753, plain, ((e1 = e4) | (~ spl24_112 | ~ spl24_115)), inference(forward_demodulation, [], [f1173, f1161])).
fof(f7750, plain, (~ spl24_48 | ~ spl24_119 | ~ spl24_122 | spl24_251), inference(avatar_split_clause, [], [f7749, f1882, f1201, f1188, f890])).
fof(f7749, plain, (~ (e2 = op(e3, e0)) | (~ spl24_119 | ~ spl24_122 | spl24_251)), inference(forward_demodulation, [], [f7738, f1190])).
fof(f7738, plain, (~ (e2 = op(op(e0, e1), e0)) | (~ spl24_122 | spl24_251)), inference(backward_demodulation, [], [f1884, f1203])).
fof(f7721, plain, (~ spl24_47 | ~ spl24_119 | spl24_205), inference(avatar_split_clause, [], [f7720, f1651, f1188, f886])).
fof(f7720, plain, (~ (e1 = op(e3, e0)) | (~ spl24_119 | spl24_205)), inference(forward_demodulation, [], [f1653, f1190])).
fof(f7689, plain, (~ spl24_122 | ~ spl24_112), inference(avatar_split_clause, [], [f7256, f1159, f1201])).
fof(f7256, plain, (~ (op(e0, e0) = e1) | ~ spl24_112), inference(forward_demodulation, [], [f433, f1161])).
fof(f7650, plain, (~ spl24_21 | ~ spl24_75 | ~ spl24_178), inference(avatar_contradiction_clause, [], [f7649])).
fof(f7649, plain, ($false | (~ spl24_21 | ~ spl24_75 | ~ spl24_178)), inference(subsumption_resolution, [], [f7648, f483])).
fof(f7648, plain, ((e0 = e2) | (~ spl24_21 | ~ spl24_75 | ~ spl24_178)), inference(forward_demodulation, [], [f7646, f779])).
fof(f779, plain, ((e0 = op(e4, e0)) | ~ spl24_21), inference(avatar_component_clause, [], [f777])).
fof(f7646, plain, ((e2 = op(e4, e0)) | (~ spl24_75 | ~ spl24_178)), inference(backward_demodulation, [], [f1519, f1005])).
fof(f7638, plain, (~ spl24_60 | ~ spl24_110), inference(avatar_split_clause, [], [f7628, f1150, f940])).
fof(f7628, plain, (~ (e4 = op(e2, e3)) | ~ spl24_110), inference(backward_demodulation, [], [f413, f1152])).
fof(f7602, plain, (spl24_47 | ~ spl24_119 | ~ spl24_205), inference(avatar_split_clause, [], [f7601, f1651, f1188, f886])).
fof(f7601, plain, ((e1 = op(e3, e0)) | (~ spl24_119 | ~ spl24_205)), inference(forward_demodulation, [], [f1652, f1190])).
fof(f7579, plain, (~ spl24_110 | ~ spl24_99 | ~ spl24_260 | spl24_269), inference(avatar_split_clause, [], [f7578, f2020, f1944, f1104, f1150])).
fof(f7578, plain, (~ (e4 = op(e0, e3)) | (~ spl24_99 | ~ spl24_260 | spl24_269)), inference(forward_demodulation, [], [f7573, f1106])).
fof(f7573, plain, (~ (e4 = op(e0, op(e1, e0))) | (~ spl24_260 | spl24_269)), inference(backward_demodulation, [], [f2022, f1945])).
fof(f7563, plain, (spl24_60 | ~ spl24_64 | ~ spl24_218), inference(avatar_split_clause, [], [f7551, f1714, f957, f940])).
fof(f7551, plain, ((e4 = op(e2, e3)) | (~ spl24_64 | ~ spl24_218)), inference(forward_demodulation, [], [f1715, f959])).
fof(f7561, plain, (~ spl24_13 | ~ spl24_53 | spl24_138), inference(avatar_contradiction_clause, [], [f7560])).
fof(f7560, plain, ($false | (~ spl24_13 | ~ spl24_53 | spl24_138)), inference(subsumption_resolution, [], [f7559, f913])).
fof(f913, plain, ((e2 = op(e2, e4)) | ~ spl24_53), inference(avatar_component_clause, [], [f911])).
fof(f7559, plain, (~ (e2 = op(e2, e4)) | (~ spl24_13 | spl24_138)), inference(forward_demodulation, [], [f1324, f745])).
fof(f745, plain, ((e2 = op(e4, e2)) | ~ spl24_13), inference(avatar_component_clause, [], [f743])).
fof(f7474, plain, (~ spl24_112 | ~ spl24_114), inference(avatar_contradiction_clause, [], [f7473])).
fof(f7473, plain, ($false | (~ spl24_112 | ~ spl24_114)), inference(subsumption_resolution, [], [f7472, f487])).
fof(f7472, plain, ((e1 = e3) | (~ spl24_112 | ~ spl24_114)), inference(forward_demodulation, [], [f1169, f1161])).
fof(f7420, plain, (~ spl24_112 | spl24_243 | ~ spl24_251), inference(avatar_contradiction_clause, [], [f7419])).
fof(f7419, plain, ($false | (~ spl24_112 | spl24_243 | ~ spl24_251)), inference(subsumption_resolution, [], [f7415, f1161])).
fof(f7415, plain, (~ (e1 = op(e0, e2)) | (spl24_243 | ~ spl24_251)), inference(backward_demodulation, [], [f1846, f1883])).
fof(f1883, plain, ((e2 = op(op(e0, op(e0, e0)), e0)) | ~ spl24_251), inference(avatar_component_clause, [], [f1882])).
fof(f7400, plain, (~ spl24_49 | ~ spl24_29), inference(avatar_split_clause, [], [f7100, f810, f894])).
fof(f7100, plain, (~ (e3 = op(e3, e0)) | ~ spl24_29), inference(forward_demodulation, [], [f465, f812])).
fof(f812, plain, ((e3 = op(e3, e4)) | ~ spl24_29), inference(avatar_component_clause, [], [f810])).
fof(f7399, plain, (~ spl24_49 | ~ spl24_99), inference(avatar_split_clause, [], [f7398, f1104, f894])).
fof(f7398, plain, (~ (e3 = op(e3, e0)) | ~ spl24_99), inference(forward_demodulation, [], [f387, f1106])).
fof(f7366, plain, (spl24_99 | ~ spl24_119 | ~ spl24_194), inference(avatar_split_clause, [], [f7362, f1597, f1188, f1104])).
fof(f7362, plain, ((e3 = op(e1, e0)) | (~ spl24_119 | ~ spl24_194)), inference(backward_demodulation, [], [f1598, f1190])).
fof(f7357, plain, (~ spl24_112 | ~ spl24_123 | spl24_260), inference(avatar_contradiction_clause, [], [f7356])).
fof(f7356, plain, ($false | (~ spl24_112 | ~ spl24_123 | spl24_260)), inference(subsumption_resolution, [], [f7350, f1161])).
fof(f7350, plain, (~ (e1 = op(e0, e2)) | (~ spl24_123 | spl24_260)), inference(backward_demodulation, [], [f1946, f1207])).
fof(f7341, plain, (spl24_99 | ~ spl24_47 | ~ spl24_162), inference(avatar_split_clause, [], [f7340, f1438, f886, f1104])).
fof(f7340, plain, ((e3 = op(e1, e0)) | (~ spl24_47 | ~ spl24_162)), inference(forward_demodulation, [], [f1440, f888])).
fof(f7338, plain, (~ spl24_57 | ~ spl24_81 | ~ spl24_169), inference(avatar_contradiction_clause, [], [f7337])).
fof(f7337, plain, ($false | (~ spl24_57 | ~ spl24_81 | ~ spl24_169)), inference(subsumption_resolution, [], [f7336, f483])).
fof(f7336, plain, ((e0 = e2) | (~ spl24_57 | ~ spl24_81 | ~ spl24_169)), inference(forward_demodulation, [], [f7335, f1031])).
fof(f7335, plain, ((e2 = op(e1, e3)) | (~ spl24_57 | ~ spl24_169)), inference(forward_demodulation, [], [f1474, f930])).
fof(f7209, plain, (~ spl24_74 | ~ spl24_108 | spl24_199), inference(avatar_split_clause, [], [f7206, f1621, f1142, f999])).
fof(f7206, plain, (~ (e3 = op(e2, e0)) | (~ spl24_108 | spl24_199)), inference(backward_demodulation, [], [f1623, f1144])).
fof(f7201, plain, (~ spl24_112 | ~ spl24_120 | ~ spl24_179 | ~ spl24_252 | spl24_269), inference(avatar_contradiction_clause, [], [f7200])).
fof(f7200, plain, ($false | (~ spl24_112 | ~ spl24_120 | ~ spl24_179 | ~ spl24_252 | spl24_269)), inference(subsumption_resolution, [], [f7184, f1194])).
fof(f7184, plain, (~ (e4 = op(e0, e1)) | (~ spl24_112 | ~ spl24_179 | ~ spl24_252 | spl24_269)), inference(backward_demodulation, [], [f7119, f1161])).
fof(f7119, plain, (~ (e4 = op(e0, op(e0, e2))) | (~ spl24_179 | ~ spl24_252 | spl24_269)), inference(forward_demodulation, [], [f7115, f1523])).
fof(f7115, plain, (~ (e4 = op(e0, op(e2, e0))) | (~ spl24_252 | spl24_269)), inference(backward_demodulation, [], [f2022, f1889])).
fof(f7164, plain, (spl24_16 | ~ spl24_120 | ~ spl24_206), inference(avatar_contradiction_clause, [], [f7163])).
fof(f7163, plain, ($false | (spl24_16 | ~ spl24_120 | ~ spl24_206)), inference(subsumption_resolution, [], [f7156, f757])).
fof(f757, plain, (~ (e0 = op(e4, e1)) | spl24_16), inference(avatar_component_clause, [], [f756])).
fof(f7156, plain, ((e0 = op(e4, e1)) | (~ spl24_120 | ~ spl24_206)), inference(backward_demodulation, [], [f1658, f1194])).
fof(f7117, plain, (~ spl24_74 | spl24_244 | ~ spl24_252), inference(avatar_split_clause, [], [f7110, f1888, f1848, f999])).
fof(f7110, plain, (~ (e3 = op(e2, e0)) | (spl24_244 | ~ spl24_252)), inference(backward_demodulation, [], [f1850, f1889])).
fof(f7088, plain, (~ spl24_118 | ~ spl24_81 | ~ spl24_94 | spl24_235), inference(avatar_split_clause, [], [f7009, f1794, f1083, f1029, f1184])).
fof(f7009, plain, (~ (e2 = op(e0, e1)) | (~ spl24_81 | ~ spl24_94 | spl24_235)), inference(forward_demodulation, [], [f7008, f1031])).
fof(f7008, plain, (~ (e2 = op(op(e1, e3), e1)) | (~ spl24_94 | spl24_235)), inference(forward_demodulation, [], [f1796, f1085])).
fof(f7058, plain, (~ spl24_115 | spl24_221 | ~ spl24_253), inference(avatar_split_clause, [], [f7057, f1893, f1728, f1171])).
fof(f7057, plain, (~ (e4 = op(e0, e2)) | (spl24_221 | ~ spl24_253)), inference(forward_demodulation, [], [f1730, f1894])).
fof(f7052, plain, (~ spl24_65 | ~ spl24_53 | spl24_165), inference(avatar_split_clause, [], [f7022, f1452, f911, f961])).
fof(f7022, plain, (~ (e4 = op(e2, e2)) | (~ spl24_53 | spl24_165)), inference(forward_demodulation, [], [f1454, f913])).
fof(f7036, plain, (~ spl24_43 | ~ spl24_225 | spl24_254), inference(avatar_split_clause, [], [f7035, f1904, f1746, f869])).
fof(f7035, plain, (~ (e2 = op(e3, e1)) | (~ spl24_225 | spl24_254)), inference(forward_demodulation, [], [f1906, f1747])).
fof(f7027, plain, (~ spl24_17 | ~ spl24_45 | ~ spl24_158), inference(avatar_contradiction_clause, [], [f7026])).
fof(f7026, plain, ($false | (~ spl24_17 | ~ spl24_45 | ~ spl24_158)), inference(subsumption_resolution, [], [f7025, f487])).
fof(f7025, plain, ((e1 = e3) | (~ spl24_17 | ~ spl24_45 | ~ spl24_158)), inference(forward_demodulation, [], [f7024, f762])).
fof(f762, plain, ((e1 = op(e4, e1)) | ~ spl24_17), inference(avatar_component_clause, [], [f760])).
fof(f7024, plain, ((e3 = op(e4, e1)) | (~ spl24_45 | ~ spl24_158)), inference(forward_demodulation, [], [f1421, f879])).
fof(f6990, plain, (~ spl24_62 | ~ spl24_66 | ~ spl24_114 | spl24_217), inference(avatar_contradiction_clause, [], [f6989])).
fof(f6989, plain, ($false | (~ spl24_62 | ~ spl24_66 | ~ spl24_114 | spl24_217)), inference(subsumption_resolution, [], [f6988, f1169])).
fof(f6988, plain, (~ (e3 = op(e0, e2)) | (~ spl24_62 | ~ spl24_66 | spl24_217)), inference(forward_demodulation, [], [f6987, f968])).
fof(f6987, plain, (~ (e3 = op(op(e2, e1), e2)) | (~ spl24_62 | spl24_217)), inference(forward_demodulation, [], [f1712, f951])).
fof(f1712, plain, (~ (e3 = op(op(e2, op(e2, e2)), e2)) | spl24_217), inference(avatar_component_clause, [], [f1710])).
fof(f6978, plain, (~ spl24_60 | ~ spl24_62 | ~ spl24_66 | ~ spl24_114 | spl24_266), inference(avatar_contradiction_clause, [], [f6977])).
fof(f6977, plain, ($false | (~ spl24_60 | ~ spl24_62 | ~ spl24_66 | ~ spl24_114 | spl24_266)), inference(subsumption_resolution, [], [f6976, f942])).
fof(f6976, plain, (~ (e4 = op(e2, e3)) | (~ spl24_62 | ~ spl24_66 | ~ spl24_114 | spl24_266)), inference(forward_demodulation, [], [f6975, f1169])).
fof(f6975, plain, (~ (e4 = op(e2, op(e0, e2))) | (~ spl24_62 | ~ spl24_66 | spl24_266)), inference(forward_demodulation, [], [f6974, f968])).
fof(f6974, plain, (~ (e4 = op(e2, op(op(e2, e1), e2))) | (~ spl24_62 | spl24_266)), inference(forward_demodulation, [], [f2002, f951])).
fof(f6898, plain, (~ spl24_53 | ~ spl24_62 | ~ spl24_165), inference(avatar_contradiction_clause, [], [f6897])).
fof(f6897, plain, ($false | (~ spl24_53 | ~ spl24_62 | ~ spl24_165)), inference(subsumption_resolution, [], [f6896, f488])).
fof(f6896, plain, ((e1 = e4) | (~ spl24_53 | ~ spl24_62 | ~ spl24_165)), inference(forward_demodulation, [], [f6890, f951])).
fof(f6890, plain, ((e4 = op(e2, e2)) | (~ spl24_53 | ~ spl24_165)), inference(backward_demodulation, [], [f1453, f913])).
fof(f6883, plain, (~ spl24_66 | ~ spl24_118 | spl24_174), inference(avatar_contradiction_clause, [], [f6882])).
fof(f6882, plain, ($false | (~ spl24_66 | ~ spl24_118 | spl24_174)), inference(subsumption_resolution, [], [f6879, f1186])).
fof(f6879, plain, (~ (e2 = op(e0, e1)) | (~ spl24_66 | spl24_174)), inference(backward_demodulation, [], [f1499, f968])).
fof(f6814, plain, (~ spl24_33 | ~ spl24_36 | ~ spl24_45 | ~ spl24_107 | spl24_267), inference(avatar_contradiction_clause, [], [f6813])).
fof(f6813, plain, ($false | (~ spl24_33 | ~ spl24_36 | ~ spl24_45 | ~ spl24_107 | spl24_267)), inference(subsumption_resolution, [], [f6805, f879])).
fof(f6805, plain, (~ (e4 = op(e3, e1)) | (~ spl24_33 | ~ spl24_36 | ~ spl24_107 | spl24_267)), inference(backward_demodulation, [], [f6674, f1140])).
fof(f6674, plain, (~ (e4 = op(e3, op(e0, e3))) | (~ spl24_33 | ~ spl24_36 | spl24_267)), inference(forward_demodulation, [], [f6673, f842])).
fof(f6673, plain, (~ (e4 = op(e3, op(op(e3, e2), e3))) | (~ spl24_33 | spl24_267)), inference(forward_demodulation, [], [f2007, f829])).
fof(f6777, plain, (~ spl24_90 | ~ spl24_118 | ~ spl24_261 | spl24_268), inference(avatar_contradiction_clause, [], [f6776])).
fof(f6776, plain, ($false | (~ spl24_90 | ~ spl24_118 | ~ spl24_261 | spl24_268)), inference(subsumption_resolution, [], [f6769, f1068])).
fof(f6769, plain, (~ (e4 = op(e1, e2)) | (~ spl24_118 | ~ spl24_261 | spl24_268)), inference(backward_demodulation, [], [f6322, f1186])).
fof(f6322, plain, (~ (e4 = op(e1, op(e0, e1))) | (~ spl24_261 | spl24_268)), inference(forward_demodulation, [], [f2012, f1950])).
fof(f6772, plain, (~ spl24_93 | ~ spl24_118), inference(avatar_split_clause, [], [f6759, f1184, f1079])).
fof(f6759, plain, (~ (e2 = op(e1, e1)) | ~ spl24_118), inference(backward_demodulation, [], [f392, f1186])).
fof(f6757, plain, (spl24_105 | ~ spl24_125 | ~ spl24_245), inference(avatar_split_clause, [], [f6751, f1852, f1213, f1129])).
fof(f6751, plain, ((e4 = op(e0, e4)) | (~ spl24_125 | ~ spl24_245)), inference(backward_demodulation, [], [f1853, f1215])).
fof(f6679, plain, (~ spl24_107 | ~ spl24_33 | ~ spl24_36 | spl24_225), inference(avatar_split_clause, [], [f6678, f1746, f840, f827, f1138])).
fof(f6678, plain, (~ (e1 = op(e0, e3)) | (~ spl24_33 | ~ spl24_36 | spl24_225)), inference(forward_demodulation, [], [f6677, f842])).
fof(f6677, plain, (~ (e1 = op(op(e3, e2), e3)) | (~ spl24_33 | spl24_225)), inference(forward_demodulation, [], [f1748, f829])).
fof(f6656, plain, (spl24_114 | ~ spl24_36 | ~ spl24_154), inference(avatar_split_clause, [], [f6556, f1400, f840, f1167])).
fof(f6556, plain, ((e3 = op(e0, e2)) | (~ spl24_36 | ~ spl24_154)), inference(backward_demodulation, [], [f1402, f842])).
fof(f6648, plain, (~ spl24_69 | ~ spl24_62 | spl24_222), inference(avatar_split_clause, [], [f6509, f1732, f949, f978])).
fof(f6509, plain, (~ (e3 = op(e2, e1)) | (~ spl24_62 | spl24_222)), inference(backward_demodulation, [], [f1734, f951])).
fof(f6537, plain, (spl24_11 | ~ spl24_51 | ~ spl24_138), inference(avatar_contradiction_clause, [], [f6536])).
fof(f6536, plain, ($false | (spl24_11 | ~ spl24_51 | ~ spl24_138)), inference(subsumption_resolution, [], [f6529, f736])).
fof(f736, plain, (~ (e0 = op(e4, e2)) | spl24_11), inference(avatar_component_clause, [], [f735])).
fof(f6529, plain, ((e0 = op(e4, e2)) | (~ spl24_51 | ~ spl24_138)), inference(backward_demodulation, [], [f1323, f905])).
fof(f6483, plain, (~ spl24_15 | ~ spl24_90), inference(avatar_split_clause, [], [f6475, f1066, f751])).
fof(f6475, plain, (~ (e4 = op(e4, e2)) | ~ spl24_90), inference(backward_demodulation, [], [f408, f1068])).
fof(f6462, plain, (~ spl24_87 | ~ spl24_97), inference(avatar_split_clause, [], [f6455, f1096, f1054])).
fof(f6455, plain, (~ (e1 = op(e1, e2)) | ~ spl24_97), inference(backward_demodulation, [], [f443, f1098])).
fof(f6461, plain, (~ spl24_22 | ~ spl24_97), inference(avatar_split_clause, [], [f6453, f1096, f781])).
fof(f6453, plain, (~ (e1 = op(e4, e0)) | ~ spl24_97), inference(backward_demodulation, [], [f388, f1098])).
fof(f6446, plain, (~ spl24_49 | ~ spl24_109 | spl24_199), inference(avatar_split_clause, [], [f6442, f1621, f1146, f894])).
fof(f6442, plain, (~ (e3 = op(e3, e0)) | (~ spl24_109 | spl24_199)), inference(backward_demodulation, [], [f1623, f1148])).
fof(f6445, plain, (~ spl24_49 | ~ spl24_109 | spl24_163), inference(avatar_split_clause, [], [f6441, f1443, f1146, f894])).
fof(f6441, plain, (~ (e3 = op(e3, e0)) | (~ spl24_109 | spl24_163)), inference(backward_demodulation, [], [f1445, f1148])).
fof(f6398, plain, (~ spl24_121 | ~ spl24_245), inference(avatar_contradiction_clause, [], [f6397])).
fof(f6397, plain, ($false | (~ spl24_121 | ~ spl24_245)), inference(subsumption_resolution, [], [f6396, f485])).
fof(f6396, plain, ((e0 = e4) | (~ spl24_121 | ~ spl24_245)), inference(forward_demodulation, [], [f6391, f1199])).
fof(f6391, plain, ((op(e0, e0) = e4) | (~ spl24_121 | ~ spl24_245)), inference(backward_demodulation, [], [f1853, f1199])).
fof(f6378, plain, (~ spl24_102 | ~ spl24_126), inference(avatar_contradiction_clause, [], [f6377])).
fof(f6377, plain, ($false | (~ spl24_102 | ~ spl24_126)), inference(subsumption_resolution, [], [f6376, f488])).
fof(f6376, plain, ((e1 = e4) | (~ spl24_102 | ~ spl24_126)), inference(forward_demodulation, [], [f6362, f1119])).
fof(f6362, plain, ((e4 = op(e0, e4)) | ~ spl24_126), inference(backward_demodulation, [], [f329, f1220])).
fof(f6346, plain, (~ spl24_8 | spl24_211 | ~ spl24_212), inference(avatar_split_clause, [], [f6345, f1688, f1684, f722])).
fof(f1684, plain, (spl24_211 <=> (e2 = op(op(e3, op(e3, e3)), e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_211])])).
fof(f6345, plain, (~ (e2 = op(e4, e3)) | (spl24_211 | ~ spl24_212)), inference(forward_demodulation, [], [f1686, f1689])).
fof(f1686, plain, (~ (e2 = op(op(e3, op(e3, e3)), e3)) | spl24_211), inference(avatar_component_clause, [], [f1684])).
fof(f6319, plain, (~ spl24_120 | spl24_231 | ~ spl24_261), inference(avatar_split_clause, [], [f6317, f1949, f1774, f1192])).
fof(f6317, plain, (~ (e4 = op(e0, e1)) | (spl24_231 | ~ spl24_261)), inference(backward_demodulation, [], [f1776, f1950])).
fof(f6280, plain, (~ spl24_71 | ~ spl24_74), inference(avatar_contradiction_clause, [], [f6279])).
fof(f6279, plain, ($false | (~ spl24_71 | ~ spl24_74)), inference(subsumption_resolution, [], [f6278, f484])).
fof(f6278, plain, ((e0 = e3) | (~ spl24_71 | ~ spl24_74)), inference(forward_demodulation, [], [f1001, f989])).
fof(f989, plain, ((e0 = op(e2, e0)) | ~ spl24_71), inference(avatar_component_clause, [], [f987])).
fof(f6277, plain, (~ spl24_111 | ~ spl24_114), inference(avatar_contradiction_clause, [], [f6276])).
fof(f6276, plain, ($false | (~ spl24_111 | ~ spl24_114)), inference(subsumption_resolution, [], [f6275, f484])).
fof(f6275, plain, ((e0 = e3) | (~ spl24_111 | ~ spl24_114)), inference(forward_demodulation, [], [f1169, f1157])).
fof(f6274, plain, (~ spl24_118 | ~ spl24_122 | spl24_252), inference(avatar_contradiction_clause, [], [f6273])).
fof(f6273, plain, ($false | (~ spl24_118 | ~ spl24_122 | spl24_252)), inference(subsumption_resolution, [], [f6267, f1186])).
fof(f6267, plain, (~ (e2 = op(e0, e1)) | (~ spl24_122 | spl24_252)), inference(backward_demodulation, [], [f1890, f1203])).
fof(f6259, plain, (~ spl24_48 | ~ spl24_71 | ~ spl24_162), inference(avatar_contradiction_clause, [], [f6258])).
fof(f6258, plain, ($false | (~ spl24_48 | ~ spl24_71 | ~ spl24_162)), inference(subsumption_resolution, [], [f6257, f484])).
fof(f6257, plain, ((e0 = e3) | (~ spl24_48 | ~ spl24_71 | ~ spl24_162)), inference(forward_demodulation, [], [f6256, f989])).
fof(f6256, plain, ((e3 = op(e2, e0)) | (~ spl24_48 | ~ spl24_162)), inference(forward_demodulation, [], [f1440, f892])).
fof(f6240, plain, (~ spl24_122 | ~ spl24_22), inference(avatar_split_clause, [], [f6239, f781, f1201])).
fof(f6239, plain, (~ (op(e0, e0) = e1) | ~ spl24_22), inference(forward_demodulation, [], [f385, f783])).
fof(f6232, plain, (~ spl24_16 | ~ spl24_45 | ~ spl24_158), inference(avatar_contradiction_clause, [], [f6231])).
fof(f6231, plain, ($false | (~ spl24_16 | ~ spl24_45 | ~ spl24_158)), inference(subsumption_resolution, [], [f6230, f484])).
fof(f6230, plain, ((e0 = e3) | (~ spl24_16 | ~ spl24_45 | ~ spl24_158)), inference(forward_demodulation, [], [f6229, f758])).
fof(f6229, plain, ((e3 = op(e4, e1)) | (~ spl24_45 | ~ spl24_158)), inference(forward_demodulation, [], [f1421, f879])).
fof(f6222, plain, (~ spl24_22 | ~ spl24_110 | ~ spl24_118 | ~ spl24_124 | spl24_257), inference(avatar_contradiction_clause, [], [f6221])).
fof(f6221, plain, ($false | (~ spl24_22 | ~ spl24_110 | ~ spl24_118 | ~ spl24_124 | spl24_257)), inference(subsumption_resolution, [], [f6220, f1186])).
fof(f6220, plain, (~ (e2 = op(e0, e1)) | (~ spl24_22 | ~ spl24_110 | ~ spl24_124 | spl24_257)), inference(forward_demodulation, [], [f6219, f783])).
fof(f6219, plain, (~ (e2 = op(e0, op(e4, e0))) | (~ spl24_110 | ~ spl24_124 | spl24_257)), inference(forward_demodulation, [], [f6218, f1152])).
fof(f6218, plain, (~ (e2 = op(e0, op(op(e0, e3), e0))) | (~ spl24_124 | spl24_257)), inference(forward_demodulation, [], [f1926, f1211])).
fof(f6208, plain, (~ spl24_8 | ~ spl24_59 | ~ spl24_133), inference(avatar_contradiction_clause, [], [f6207])).
fof(f6207, plain, ($false | (~ spl24_8 | ~ spl24_59 | ~ spl24_133)), inference(subsumption_resolution, [], [f6206, f491])).
fof(f6206, plain, ((e3 = e4) | (~ spl24_8 | ~ spl24_59 | ~ spl24_133)), inference(forward_demodulation, [], [f6202, f938])).
fof(f6202, plain, ((e4 = op(e2, e3)) | (~ spl24_8 | ~ spl24_133)), inference(backward_demodulation, [], [f1300, f724])).
fof(f6196, plain, (~ spl24_22 | ~ spl24_100 | spl24_145), inference(avatar_contradiction_clause, [], [f6195])).
fof(f6195, plain, ($false | (~ spl24_22 | ~ spl24_100 | spl24_145)), inference(subsumption_resolution, [], [f6190, f1110])).
fof(f6190, plain, (~ (e4 = op(e1, e0)) | (~ spl24_22 | spl24_145)), inference(backward_demodulation, [], [f1356, f783])).
fof(f1356, plain, (~ (e4 = op(op(e4, e0), e0)) | spl24_145), inference(avatar_component_clause, [], [f1355])).
fof(f6194, plain, (~ spl24_7 | ~ spl24_22), inference(avatar_split_clause, [], [f6185, f781, f718])).
fof(f6185, plain, (~ (e1 = op(e4, e3)) | ~ spl24_22), inference(backward_demodulation, [], [f474, f783])).
fof(f6162, plain, (~ spl24_33 | ~ spl24_48), inference(avatar_split_clause, [], [f6156, f890, f827])).
fof(f6156, plain, (~ (e2 = op(e3, e3)) | ~ spl24_48), inference(backward_demodulation, [], [f464, f892])).
fof(f6161, plain, (~ spl24_23 | ~ spl24_48), inference(avatar_split_clause, [], [f6155, f890, f785])).
fof(f6155, plain, (~ (e2 = op(e4, e0)) | ~ spl24_48), inference(backward_demodulation, [], [f391, f892])).
fof(f6144, plain, (~ spl24_50 | ~ spl24_100), inference(avatar_split_clause, [], [f6143, f1108, f898])).
fof(f6143, plain, (~ (e4 = op(e3, e0)) | ~ spl24_100), inference(forward_demodulation, [], [f387, f1110])).
fof(f6085, plain, (~ spl24_110 | ~ spl24_124 | spl24_245), inference(avatar_contradiction_clause, [], [f6084])).
fof(f6084, plain, ($false | (~ spl24_110 | ~ spl24_124 | spl24_245)), inference(subsumption_resolution, [], [f6083, f1152])).
fof(f6083, plain, (~ (e4 = op(e0, e3)) | (~ spl24_124 | spl24_245)), inference(forward_demodulation, [], [f1854, f1211])).
fof(f6026, plain, (spl24_66 | ~ spl24_118 | ~ spl24_206), inference(avatar_contradiction_clause, [], [f6025])).
fof(f6025, plain, ($false | (spl24_66 | ~ spl24_118 | ~ spl24_206)), inference(subsumption_resolution, [], [f6019, f967])).
fof(f6019, plain, ((e0 = op(e2, e1)) | (~ spl24_118 | ~ spl24_206)), inference(backward_demodulation, [], [f1658, f1186])).
fof(f5996, plain, (~ spl24_110 | ~ spl24_26 | spl24_148), inference(avatar_split_clause, [], [f5755, f1369, f798, f1150])).
fof(f5755, plain, (~ (e4 = op(e0, e3)) | (~ spl24_26 | spl24_148)), inference(forward_demodulation, [], [f1371, f800])).
fof(f5987, plain, (~ spl24_110 | spl24_223 | ~ spl24_249), inference(avatar_split_clause, [], [f5986, f1871, f1737, f1150])).
fof(f5986, plain, (~ (e4 = op(e0, e3)) | (spl24_223 | ~ spl24_249)), inference(forward_demodulation, [], [f1739, f1872])).
fof(f5974, plain, (~ spl24_23 | ~ spl24_71 | ~ spl24_145), inference(avatar_contradiction_clause, [], [f5973])).
fof(f5973, plain, ($false | (~ spl24_23 | ~ spl24_71 | ~ spl24_145)), inference(subsumption_resolution, [], [f5972, f485])).
fof(f5972, plain, ((e0 = e4) | (~ spl24_23 | ~ spl24_71 | ~ spl24_145)), inference(forward_demodulation, [], [f5971, f989])).
fof(f5971, plain, ((e4 = op(e2, e0)) | (~ spl24_23 | ~ spl24_145)), inference(forward_demodulation, [], [f1357, f787])).
fof(f5959, plain, (~ spl24_81 | ~ spl24_94 | spl24_261), inference(avatar_contradiction_clause, [], [f5958])).
fof(f5958, plain, ($false | (~ spl24_81 | ~ spl24_94 | spl24_261)), inference(subsumption_resolution, [], [f5957, f1031])).
fof(f5957, plain, (~ (e0 = op(e1, e3)) | (~ spl24_94 | spl24_261)), inference(forward_demodulation, [], [f1951, f1085])).
fof(f5946, plain, (~ spl24_78 | ~ spl24_81 | ~ spl24_94 | ~ spl24_120 | spl24_256), inference(avatar_contradiction_clause, [], [f5945])).
fof(f5945, plain, ($false | (~ spl24_78 | ~ spl24_81 | ~ spl24_94 | ~ spl24_120 | spl24_256)), inference(subsumption_resolution, [], [f5944, f1018])).
fof(f5944, plain, (~ (e2 = op(e1, e4)) | (~ spl24_81 | ~ spl24_94 | ~ spl24_120 | spl24_256)), inference(forward_demodulation, [], [f5943, f1194])).
fof(f5943, plain, (~ (e2 = op(e1, op(e0, e1))) | (~ spl24_81 | ~ spl24_94 | spl24_256)), inference(forward_demodulation, [], [f5933, f1031])).
fof(f5933, plain, (~ (e2 = op(e1, op(op(e1, e3), e1))) | (~ spl24_94 | spl24_256)), inference(backward_demodulation, [], [f1916, f1085])).
fof(f5942, plain, (~ spl24_81 | ~ spl24_94 | ~ spl24_120 | spl24_231), inference(avatar_contradiction_clause, [], [f5941])).
fof(f5941, plain, ($false | (~ spl24_81 | ~ spl24_94 | ~ spl24_120 | spl24_231)), inference(subsumption_resolution, [], [f5940, f1194])).
fof(f5940, plain, (~ (e4 = op(e0, e1)) | (~ spl24_81 | ~ spl24_94 | spl24_231)), inference(forward_demodulation, [], [f5932, f1031])).
fof(f5932, plain, (~ (e4 = op(op(e1, e3), e1)) | (~ spl24_94 | spl24_231)), inference(backward_demodulation, [], [f1776, f1085])).
fof(f5928, plain, (spl24_22 | ~ spl24_100 | ~ spl24_193), inference(avatar_contradiction_clause, [], [f5927])).
fof(f5927, plain, ($false | (spl24_22 | ~ spl24_100 | ~ spl24_193)), inference(subsumption_resolution, [], [f5923, f782])).
fof(f782, plain, (~ (e1 = op(e4, e0)) | spl24_22), inference(avatar_component_clause, [], [f781])).
fof(f5923, plain, ((e1 = op(e4, e0)) | (~ spl24_100 | ~ spl24_193)), inference(backward_demodulation, [], [f1594, f1110])).
fof(f5926, plain, (~ spl24_16 | ~ spl24_100 | spl24_192), inference(avatar_contradiction_clause, [], [f5925])).
fof(f5925, plain, ($false | (~ spl24_16 | ~ spl24_100 | spl24_192)), inference(subsumption_resolution, [], [f5922, f758])).
fof(f5922, plain, (~ (e0 = op(e4, e1)) | (~ spl24_100 | spl24_192)), inference(backward_demodulation, [], [f1589, f1110])).
fof(f5920, plain, (~ spl24_16 | ~ spl24_120 | spl24_206), inference(avatar_contradiction_clause, [], [f5919])).
fof(f5919, plain, ($false | (~ spl24_16 | ~ spl24_120 | spl24_206)), inference(subsumption_resolution, [], [f5914, f758])).
fof(f5914, plain, (~ (e0 = op(e4, e1)) | (~ spl24_120 | spl24_206)), inference(backward_demodulation, [], [f1657, f1194])).
fof(f5915, plain, (~ spl24_95 | ~ spl24_120), inference(avatar_split_clause, [], [f5911, f1192, f1087])).
fof(f5911, plain, (~ (e4 = op(e1, e1)) | ~ spl24_120), inference(backward_demodulation, [], [f392, f1194])).
fof(f5854, plain, (~ spl24_3 | ~ spl24_4), inference(avatar_contradiction_clause, [], [f5853])).
fof(f5853, plain, ($false | (~ spl24_3 | ~ spl24_4)), inference(subsumption_resolution, [], [f5852, f489])).
fof(f5852, plain, ((e2 = e3) | (~ spl24_3 | ~ spl24_4)), inference(forward_demodulation, [], [f707, f703])).
fof(f5836, plain, (~ spl24_99 | ~ spl24_102 | ~ spl24_108 | ~ spl24_125 | spl24_257), inference(avatar_contradiction_clause, [], [f5835])).
fof(f5835, plain, ($false | (~ spl24_99 | ~ spl24_102 | ~ spl24_108 | ~ spl24_125 | spl24_257)), inference(subsumption_resolution, [], [f5830, f1144])).
fof(f5830, plain, (~ (e2 = op(e0, e3)) | (~ spl24_99 | ~ spl24_102 | ~ spl24_125 | spl24_257)), inference(backward_demodulation, [], [f5694, f1106])).
fof(f5694, plain, (~ (e2 = op(e0, op(e1, e0))) | (~ spl24_102 | ~ spl24_125 | spl24_257)), inference(backward_demodulation, [], [f5605, f1119])).
fof(f5605, plain, (~ (e2 = op(e0, op(op(e0, e4), e0))) | (~ spl24_125 | spl24_257)), inference(forward_demodulation, [], [f1926, f1215])).
fof(f5798, plain, (~ spl24_71 | ~ spl24_111 | spl24_179), inference(avatar_contradiction_clause, [], [f5797])).
fof(f5797, plain, ($false | (~ spl24_71 | ~ spl24_111 | spl24_179)), inference(subsumption_resolution, [], [f5796, f1157])).
fof(f5796, plain, (~ (e0 = op(e0, e2)) | (~ spl24_71 | spl24_179)), inference(forward_demodulation, [], [f1524, f989])).
fof(f5784, plain, (~ spl24_71 | ~ spl24_98 | ~ spl24_193), inference(avatar_contradiction_clause, [], [f5783])).
fof(f5783, plain, ($false | (~ spl24_71 | ~ spl24_98 | ~ spl24_193)), inference(subsumption_resolution, [], [f5782, f482])).
fof(f5782, plain, ((e0 = e1) | (~ spl24_71 | ~ spl24_98 | ~ spl24_193)), inference(forward_demodulation, [], [f5781, f989])).
fof(f5781, plain, ((e1 = op(e2, e0)) | (~ spl24_98 | ~ spl24_193)), inference(forward_demodulation, [], [f1594, f1102])).
fof(f5780, plain, (~ spl24_43 | ~ spl24_67 | ~ spl24_158), inference(avatar_contradiction_clause, [], [f5779])).
fof(f5779, plain, ($false | (~ spl24_43 | ~ spl24_67 | ~ spl24_158)), inference(subsumption_resolution, [], [f5778, f487])).
fof(f5778, plain, ((e1 = e3) | (~ spl24_43 | ~ spl24_67 | ~ spl24_158)), inference(forward_demodulation, [], [f5777, f972])).
fof(f972, plain, ((e1 = op(e2, e1)) | ~ spl24_67), inference(avatar_component_clause, [], [f970])).
fof(f5777, plain, ((e3 = op(e2, e1)) | (~ spl24_43 | ~ spl24_158)), inference(forward_demodulation, [], [f1421, f871])).
fof(f5769, plain, (~ spl24_81 | spl24_107 | ~ spl24_185), inference(avatar_contradiction_clause, [], [f5768])).
fof(f5768, plain, ($false | (~ spl24_81 | spl24_107 | ~ spl24_185)), inference(subsumption_resolution, [], [f5767, f1139])).
fof(f1139, plain, (~ (e1 = op(e0, e3)) | spl24_107), inference(avatar_component_clause, [], [f1138])).
fof(f5767, plain, ((e1 = op(e0, e3)) | (~ spl24_81 | ~ spl24_185)), inference(forward_demodulation, [], [f1553, f1031])).
fof(f5762, plain, (~ spl24_15 | ~ spl24_55 | spl24_138), inference(avatar_contradiction_clause, [], [f5761])).
fof(f5761, plain, ($false | (~ spl24_15 | ~ spl24_55 | spl24_138)), inference(subsumption_resolution, [], [f5760, f921])).
fof(f5760, plain, (~ (e4 = op(e2, e4)) | (~ spl24_15 | spl24_138)), inference(forward_demodulation, [], [f1324, f753])).
fof(f753, plain, ((e4 = op(e4, e2)) | ~ spl24_15), inference(avatar_component_clause, [], [f751])).
fof(f5665, plain, (~ spl24_67 | ~ spl24_87 | spl24_173), inference(avatar_contradiction_clause, [], [f5664])).
fof(f5664, plain, ($false | (~ spl24_67 | ~ spl24_87 | spl24_173)), inference(subsumption_resolution, [], [f5663, f1056])).
fof(f5663, plain, (~ (e1 = op(e1, e2)) | (~ spl24_67 | spl24_173)), inference(forward_demodulation, [], [f1495, f972])).
fof(f5659, plain, (~ spl24_67 | ~ spl24_87 | spl24_175), inference(avatar_contradiction_clause, [], [f5658])).
fof(f5658, plain, ($false | (~ spl24_67 | ~ spl24_87 | spl24_175)), inference(subsumption_resolution, [], [f5657, f1056])).
fof(f5657, plain, (~ (e1 = op(e1, e2)) | (~ spl24_67 | spl24_175)), inference(forward_demodulation, [], [f1505, f972])).
fof(f5638, plain, (~ spl24_102 | ~ spl24_16 | spl24_140), inference(avatar_split_clause, [], [f5569, f1331, f756, f1117])).
fof(f5569, plain, (~ (e1 = op(e0, e4)) | (~ spl24_16 | spl24_140)), inference(forward_demodulation, [], [f1333, f758])).
fof(f5617, plain, (spl24_41 | ~ spl24_119 | ~ spl24_206), inference(avatar_contradiction_clause, [], [f5616])).
fof(f5616, plain, ($false | (spl24_41 | ~ spl24_119 | ~ spl24_206)), inference(subsumption_resolution, [], [f5615, f862])).
fof(f862, plain, (~ (e0 = op(e3, e1)) | spl24_41), inference(avatar_component_clause, [], [f861])).
fof(f5615, plain, ((e0 = op(e3, e1)) | (~ spl24_119 | ~ spl24_206)), inference(forward_demodulation, [], [f1658, f1190])).
fof(f5612, plain, (~ spl24_26 | ~ spl24_35 | ~ spl24_43 | ~ spl24_107 | spl24_254), inference(avatar_contradiction_clause, [], [f5611])).
fof(f5611, plain, ($false | (~ spl24_26 | ~ spl24_35 | ~ spl24_43 | ~ spl24_107 | spl24_254)), inference(subsumption_resolution, [], [f5610, f871])).
fof(f5610, plain, (~ (e2 = op(e3, e1)) | (~ spl24_26 | ~ spl24_35 | ~ spl24_107 | spl24_254)), inference(forward_demodulation, [], [f5609, f1140])).
fof(f5609, plain, (~ (e2 = op(e3, op(e0, e3))) | (~ spl24_26 | ~ spl24_35 | spl24_254)), inference(forward_demodulation, [], [f5608, f800])).
fof(f5608, plain, (~ (e2 = op(e3, op(op(e3, e4), e3))) | (~ spl24_35 | spl24_254)), inference(forward_demodulation, [], [f1906, f837])).
fof(f5574, plain, (~ spl24_104 | ~ spl24_119), inference(avatar_split_clause, [], [f5571, f1188, f1125])).
fof(f5571, plain, (~ (e3 = op(e0, e4)) | ~ spl24_119), inference(backward_demodulation, [], [f438, f1190])).
fof(f5573, plain, (~ spl24_44 | ~ spl24_119), inference(avatar_split_clause, [], [f5570, f1188, f873])).
fof(f5570, plain, (~ (e3 = op(e3, e1)) | ~ spl24_119), inference(backward_demodulation, [], [f394, f1190])).
fof(f5568, plain, (~ spl24_16 | spl24_120 | ~ spl24_141), inference(avatar_contradiction_clause, [], [f5567])).
fof(f5567, plain, ($false | (~ spl24_16 | spl24_120 | ~ spl24_141)), inference(subsumption_resolution, [], [f5566, f1193])).
fof(f1193, plain, (~ (e4 = op(e0, e1)) | spl24_120), inference(avatar_component_clause, [], [f1192])).
fof(f5566, plain, ((e4 = op(e0, e1)) | (~ spl24_16 | ~ spl24_141)), inference(forward_demodulation, [], [f1338, f758])).
fof(f5536, plain, (~ spl24_2 | ~ spl24_3), inference(avatar_contradiction_clause, [], [f5535])).
fof(f5535, plain, ($false | (~ spl24_2 | ~ spl24_3)), inference(subsumption_resolution, [], [f5534, f486])).
fof(f486, plain, ~ (e1 = e2), inference(cnf_transformation, [], [f5])).
fof(f5534, plain, ((e1 = e2) | (~ spl24_2 | ~ spl24_3)), inference(forward_demodulation, [], [f703, f699])).
fof(f5526, plain, (~ spl24_38 | ~ spl24_39), inference(avatar_contradiction_clause, [], [f5525])).
fof(f5525, plain, ($false | (~ spl24_38 | ~ spl24_39)), inference(subsumption_resolution, [], [f5524, f489])).
fof(f5524, plain, ((e2 = e3) | (~ spl24_38 | ~ spl24_39)), inference(backward_demodulation, [], [f854, f850])).
fof(f854, plain, ((e3 = op(e3, e2)) | ~ spl24_39), inference(avatar_component_clause, [], [f852])).
fof(f5480, plain, (spl24_102 | ~ spl24_125 | ~ spl24_260), inference(avatar_contradiction_clause, [], [f5479])).
fof(f5479, plain, ($false | (spl24_102 | ~ spl24_125 | ~ spl24_260)), inference(subsumption_resolution, [], [f5476, f1118])).
fof(f1118, plain, (~ (e1 = op(e0, e4)) | spl24_102), inference(avatar_component_clause, [], [f1117])).
fof(f5476, plain, ((e1 = op(e0, e4)) | (~ spl24_125 | ~ spl24_260)), inference(backward_demodulation, [], [f1945, f1215])).
fof(f5455, plain, (~ spl24_26 | ~ spl24_35 | ~ spl24_107 | spl24_225), inference(avatar_contradiction_clause, [], [f5454])).
fof(f5454, plain, ($false | (~ spl24_26 | ~ spl24_35 | ~ spl24_107 | spl24_225)), inference(subsumption_resolution, [], [f5453, f1140])).
fof(f5453, plain, (~ (e1 = op(e0, e3)) | (~ spl24_26 | ~ spl24_35 | spl24_225)), inference(forward_demodulation, [], [f5452, f800])).
fof(f5452, plain, (~ (e1 = op(op(e3, e4), e3)) | (~ spl24_35 | spl24_225)), inference(forward_demodulation, [], [f1748, f837])).
fof(f5433, plain, (~ spl24_100 | spl24_247 | ~ spl24_260), inference(avatar_split_clause, [], [f5430, f1944, f1862, f1108])).
fof(f5430, plain, (~ (e4 = op(e1, e0)) | (spl24_247 | ~ spl24_260)), inference(backward_demodulation, [], [f1864, f1945])).
fof(f5421, plain, (~ spl24_44 | ~ spl24_39), inference(avatar_split_clause, [], [f5293, f852, f873])).
fof(f5293, plain, (~ (e3 = op(e3, e1)) | ~ spl24_39), inference(forward_demodulation, [], [f466, f854])).
fof(f5418, plain, (~ spl24_16 | ~ spl24_19), inference(avatar_contradiction_clause, [], [f5417])).
fof(f5417, plain, ($false | (~ spl24_16 | ~ spl24_19)), inference(subsumption_resolution, [], [f5416, f484])).
fof(f5416, plain, ((e0 = e3) | (~ spl24_16 | ~ spl24_19)), inference(forward_demodulation, [], [f770, f758])).
fof(f5369, plain, (spl24_51 | ~ spl24_103 | ~ spl24_197), inference(avatar_contradiction_clause, [], [f5368])).
fof(f5368, plain, ($false | (spl24_51 | ~ spl24_103 | ~ spl24_197)), inference(subsumption_resolution, [], [f5367, f904])).
fof(f904, plain, (~ (e0 = op(e2, e4)) | spl24_51), inference(avatar_component_clause, [], [f903])).
fof(f5367, plain, ((e0 = op(e2, e4)) | (~ spl24_103 | ~ spl24_197)), inference(forward_demodulation, [], [f1613, f1123])).
fof(f5304, plain, (~ spl24_22 | ~ spl24_47), inference(avatar_split_clause, [], [f5303, f886, f781])).
fof(f5303, plain, (~ (e1 = op(e4, e0)) | ~ spl24_47), inference(forward_demodulation, [], [f391, f888])).
fof(f5265, plain, (~ spl24_2 | ~ spl24_4), inference(avatar_contradiction_clause, [], [f5264])).
fof(f5264, plain, ($false | (~ spl24_2 | ~ spl24_4)), inference(subsumption_resolution, [], [f5263, f487])).
fof(f5263, plain, ((e1 = e3) | (~ spl24_2 | ~ spl24_4)), inference(backward_demodulation, [], [f707, f699])).
fof(f5244, plain, (~ spl24_47 | ~ spl24_48), inference(avatar_contradiction_clause, [], [f5243])).
fof(f5243, plain, ($false | (~ spl24_47 | ~ spl24_48)), inference(subsumption_resolution, [], [f5242, f486])).
fof(f5242, plain, ((e1 = e2) | (~ spl24_47 | ~ spl24_48)), inference(forward_demodulation, [], [f892, f888])).
fof(f5241, plain, (~ spl24_15 | ~ spl24_55 | spl24_165), inference(avatar_split_clause, [], [f5240, f1452, f919, f751])).
fof(f5240, plain, (~ (e4 = op(e4, e2)) | (~ spl24_55 | spl24_165)), inference(backward_demodulation, [], [f1454, f921])).
fof(f5224, plain, (~ spl24_26 | ~ spl24_79 | ~ spl24_182), inference(avatar_contradiction_clause, [], [f5223])).
fof(f5223, plain, ($false | (~ spl24_26 | ~ spl24_79 | ~ spl24_182)), inference(subsumption_resolution, [], [f5222, f482])).
fof(f5222, plain, ((e0 = e1) | (~ spl24_26 | ~ spl24_79 | ~ spl24_182)), inference(forward_demodulation, [], [f5221, f800])).
fof(f5221, plain, ((e1 = op(e3, e4)) | (~ spl24_79 | ~ spl24_182)), inference(backward_demodulation, [], [f1538, f1022])).
fof(f5216, plain, (~ spl24_81 | ~ spl24_85), inference(avatar_contradiction_clause, [], [f5215])).
fof(f5215, plain, ($false | (~ spl24_81 | ~ spl24_85)), inference(subsumption_resolution, [], [f5214, f485])).
fof(f5214, plain, ((e0 = e4) | (~ spl24_81 | ~ spl24_85)), inference(forward_demodulation, [], [f1047, f1031])).
fof(f5213, plain, (~ spl24_87 | ~ spl24_88), inference(avatar_contradiction_clause, [], [f5212])).
fof(f5212, plain, ($false | (~ spl24_87 | ~ spl24_88)), inference(subsumption_resolution, [], [f5211, f486])).
fof(f5211, plain, ((e1 = e2) | (~ spl24_87 | ~ spl24_88)), inference(forward_demodulation, [], [f1060, f1056])).
fof(f5197, plain, (~ spl24_100 | ~ spl24_103 | ~ spl24_107 | ~ spl24_124 | spl24_257), inference(avatar_contradiction_clause, [], [f5196])).
fof(f5196, plain, ($false | (~ spl24_100 | ~ spl24_103 | ~ spl24_107 | ~ spl24_124 | spl24_257)), inference(subsumption_resolution, [], [f5189, f1123])).
fof(f5189, plain, (~ (e2 = op(e0, e4)) | (~ spl24_100 | ~ spl24_107 | ~ spl24_124 | spl24_257)), inference(backward_demodulation, [], [f4705, f1110])).
fof(f4705, plain, (~ (e2 = op(e0, op(e1, e0))) | (~ spl24_107 | ~ spl24_124 | spl24_257)), inference(forward_demodulation, [], [f4704, f1140])).
fof(f4704, plain, (~ (e2 = op(e0, op(op(e0, e3), e0))) | (~ spl24_124 | spl24_257)), inference(forward_demodulation, [], [f1926, f1211])).
fof(f5191, plain, (~ spl24_80 | ~ spl24_100), inference(avatar_split_clause, [], [f5186, f1108, f1024])).
fof(f5186, plain, (~ (e4 = op(e1, e4)) | ~ spl24_100), inference(backward_demodulation, [], [f445, f1110])).
fof(f5150, plain, (~ spl24_98 | ~ spl24_107 | ~ spl24_124 | spl24_251), inference(avatar_split_clause, [], [f5149, f1882, f1209, f1138, f1100])).
fof(f5149, plain, (~ (e2 = op(e1, e0)) | (~ spl24_107 | ~ spl24_124 | spl24_251)), inference(forward_demodulation, [], [f5148, f1140])).
fof(f5148, plain, (~ (e2 = op(op(e0, e3), e0)) | (~ spl24_124 | spl24_251)), inference(forward_demodulation, [], [f1884, f1211])).
fof(f5114, plain, (~ spl24_92 | ~ spl24_67), inference(avatar_split_clause, [], [f5113, f970, f1075])).
fof(f5113, plain, (~ (e1 = op(e1, e1)) | ~ spl24_67), inference(forward_demodulation, [], [f396, f972])).
fof(f5112, plain, (~ spl24_92 | ~ spl24_87), inference(avatar_split_clause, [], [f5111, f1054, f1075])).
fof(f5111, plain, (~ (e1 = op(e1, e1)) | ~ spl24_87), inference(forward_demodulation, [], [f446, f1056])).
fof(f5110, plain, (~ spl24_93 | ~ spl24_87 | spl24_187), inference(avatar_split_clause, [], [f4966, f1561, f1054, f1079])).
fof(f4966, plain, (~ (e2 = op(e1, e1)) | (~ spl24_87 | spl24_187)), inference(backward_demodulation, [], [f1563, f1056])).
fof(f5106, plain, (~ spl24_79 | ~ spl24_231 | spl24_264), inference(avatar_split_clause, [], [f4860, f1970, f1774, f1020])).
fof(f4860, plain, (~ (e3 = op(e1, e4)) | (~ spl24_231 | spl24_264)), inference(forward_demodulation, [], [f1972, f1775])).
fof(f1775, plain, ((e4 = op(op(e1, op(e1, e1)), e1)) | ~ spl24_231), inference(avatar_component_clause, [], [f1774])).
fof(f5084, plain, (~ spl24_39 | ~ spl24_59 | spl24_155), inference(avatar_contradiction_clause, [], [f5083])).
fof(f5083, plain, ($false | (~ spl24_39 | ~ spl24_59 | spl24_155)), inference(subsumption_resolution, [], [f5082, f938])).
fof(f5082, plain, (~ (e3 = op(e2, e3)) | (~ spl24_39 | spl24_155)), inference(forward_demodulation, [], [f1407, f854])).
fof(f5074, plain, (~ spl24_54 | ~ spl24_59), inference(avatar_split_clause, [], [f5073, f936, f915])).
fof(f5073, plain, (~ (e3 = op(e2, e4)) | ~ spl24_59), inference(forward_demodulation, [], [f461, f938])).
fof(f5068, plain, (~ spl24_34 | ~ spl24_39), inference(avatar_split_clause, [], [f5067, f852, f831])).
fof(f5067, plain, (~ (e3 = op(e3, e3)) | ~ spl24_39), inference(forward_demodulation, [], [f469, f854])).
fof(f5025, plain, (~ spl24_29 | ~ spl24_39), inference(avatar_split_clause, [], [f5021, f852, f810])).
fof(f5021, plain, (~ (e3 = op(e3, e4)) | ~ spl24_39), inference(backward_demodulation, [], [f470, f854])).
fof(f5016, plain, (~ spl24_47 | ~ spl24_81 | spl24_161), inference(avatar_contradiction_clause, [], [f5015])).
fof(f5015, plain, ($false | (~ spl24_47 | ~ spl24_81 | spl24_161)), inference(subsumption_resolution, [], [f5012, f1031])).
fof(f5012, plain, (~ (e0 = op(e1, e3)) | (~ spl24_47 | spl24_161)), inference(backward_demodulation, [], [f1435, f888])).
fof(f4988, plain, (~ spl24_16 | ~ spl24_80 | ~ spl24_181), inference(avatar_contradiction_clause, [], [f4987])).
fof(f4987, plain, ($false | (~ spl24_16 | ~ spl24_80 | ~ spl24_181)), inference(subsumption_resolution, [], [f4986, f485])).
fof(f4986, plain, ((e0 = e4) | (~ spl24_16 | ~ spl24_80 | ~ spl24_181)), inference(forward_demodulation, [], [f4983, f758])).
fof(f4983, plain, ((e4 = op(e4, e1)) | (~ spl24_80 | ~ spl24_181)), inference(backward_demodulation, [], [f1532, f1026])).
fof(f1532, plain, ((e4 = op(op(e1, e4), e1)) | ~ spl24_181), inference(avatar_component_clause, [], [f1531])).
fof(f4980, plain, (~ spl24_81 | ~ spl24_84), inference(avatar_contradiction_clause, [], [f4979])).
fof(f4979, plain, ($false | (~ spl24_81 | ~ spl24_84)), inference(subsumption_resolution, [], [f4978, f484])).
fof(f4978, plain, ((e0 = e3) | (~ spl24_81 | ~ spl24_84)), inference(forward_demodulation, [], [f1043, f1031])).
fof(f4922, plain, (~ spl24_111 | ~ spl24_115), inference(avatar_contradiction_clause, [], [f4921])).
fof(f4921, plain, ($false | (~ spl24_111 | ~ spl24_115)), inference(subsumption_resolution, [], [f4920, f485])).
fof(f4920, plain, ((e0 = e4) | (~ spl24_111 | ~ spl24_115)), inference(backward_demodulation, [], [f1173, f1157])).
fof(f4916, plain, (~ spl24_13 | ~ spl24_128), inference(avatar_contradiction_clause, [], [f4915])).
fof(f4915, plain, ($false | (~ spl24_13 | ~ spl24_128)), inference(subsumption_resolution, [], [f4914, f490])).
fof(f4914, plain, ((e2 = e4) | (~ spl24_13 | ~ spl24_128)), inference(forward_demodulation, [], [f4898, f745])).
fof(f4898, plain, ((e4 = op(e4, e2)) | ~ spl24_128), inference(backward_demodulation, [], [f330, f1228])).
fof(f4913, plain, (~ spl24_53 | ~ spl24_128), inference(avatar_contradiction_clause, [], [f4912])).
fof(f4912, plain, ($false | (~ spl24_53 | ~ spl24_128)), inference(subsumption_resolution, [], [f4911, f490])).
fof(f4911, plain, ((e2 = e4) | (~ spl24_53 | ~ spl24_128)), inference(forward_demodulation, [], [f4897, f913])).
fof(f4897, plain, ((e4 = op(e2, e4)) | ~ spl24_128), inference(backward_demodulation, [], [f329, f1228])).
fof(f4825, plain, (~ spl24_47 | ~ spl24_72), inference(avatar_split_clause, [], [f4677, f991, f886])).
fof(f4677, plain, (~ (e1 = op(e3, e0)) | ~ spl24_72), inference(forward_demodulation, [], [f389, f993])).
fof(f4799, plain, (~ spl24_26 | ~ spl24_16 | ~ spl24_134), inference(avatar_split_clause, [], [f4798, f1303, f756, f798])).
fof(f4798, plain, (~ (e0 = op(e3, e4)) | (~ spl24_16 | ~ spl24_134)), inference(forward_demodulation, [], [f4797, f758])).
fof(f4797, plain, (~ (op(e3, e4) = op(e4, e1)) | ~ spl24_134), inference(forward_demodulation, [], [f477, f1304])).
fof(f4787, plain, (~ spl24_16 | ~ spl24_17), inference(avatar_contradiction_clause, [], [f4786])).
fof(f4786, plain, ($false | (~ spl24_16 | ~ spl24_17)), inference(subsumption_resolution, [], [f4785, f482])).
fof(f4785, plain, ((e0 = e1) | (~ spl24_16 | ~ spl24_17)), inference(backward_demodulation, [], [f762, f758])).
fof(f4772, plain, (~ spl24_32 | ~ spl24_37), inference(avatar_split_clause, [], [f4766, f844, f823])).
fof(f4766, plain, (~ (e1 = op(e3, e3)) | ~ spl24_37), inference(backward_demodulation, [], [f469, f846])).
fof(f4755, plain, (~ spl24_21 | ~ spl24_46), inference(avatar_split_clause, [], [f4745, f882, f777])).
fof(f4745, plain, (~ (e0 = op(e4, e0)) | ~ spl24_46), inference(backward_demodulation, [], [f391, f884])).
fof(f4742, plain, (~ spl24_98 | ~ spl24_100), inference(avatar_contradiction_clause, [], [f4741])).
fof(f4741, plain, ($false | (~ spl24_98 | ~ spl24_100)), inference(subsumption_resolution, [], [f4740, f490])).
fof(f4740, plain, ((e2 = e4) | (~ spl24_98 | ~ spl24_100)), inference(forward_demodulation, [], [f1110, f1102])).
fof(f4665, plain, (~ spl24_42 | ~ spl24_17), inference(avatar_split_clause, [], [f4664, f760, f865])).
fof(f4664, plain, (~ (e1 = op(e3, e1)) | ~ spl24_17), inference(forward_demodulation, [], [f401, f762])).
fof(f4663, plain, (spl24_41 | ~ spl24_81 | ~ spl24_159), inference(avatar_split_clause, [], [f4662, f1424, f1029, f861])).
fof(f4662, plain, ((e0 = op(e3, e1)) | (~ spl24_81 | ~ spl24_159)), inference(forward_demodulation, [], [f1425, f1031])).
fof(f4652, plain, (~ spl24_23 | ~ spl24_13), inference(avatar_split_clause, [], [f4651, f743, f785])).
fof(f4651, plain, (~ (e2 = op(e4, e0)) | ~ spl24_13), inference(forward_demodulation, [], [f473, f745])).
fof(f4650, plain, (~ spl24_22 | ~ spl24_17), inference(avatar_split_clause, [], [f4649, f760, f781])).
fof(f4649, plain, (~ (e1 = op(e4, e0)) | ~ spl24_17), inference(forward_demodulation, [], [f472, f762])).
fof(f4629, plain, (~ spl24_61 | ~ spl24_62), inference(avatar_contradiction_clause, [], [f4628])).
fof(f4628, plain, ($false | (~ spl24_61 | ~ spl24_62)), inference(subsumption_resolution, [], [f4627, f482])).
fof(f4627, plain, ((e0 = e1) | (~ spl24_61 | ~ spl24_62)), inference(forward_demodulation, [], [f951, f947])).
fof(f4618, plain, (~ spl24_81 | ~ spl24_83), inference(avatar_contradiction_clause, [], [f4617])).
fof(f4617, plain, ($false | (~ spl24_81 | ~ spl24_83)), inference(subsumption_resolution, [], [f4616, f483])).
fof(f4616, plain, ((e0 = e2) | (~ spl24_81 | ~ spl24_83)), inference(forward_demodulation, [], [f1039, f1031])).
fof(f4575, plain, (~ spl24_81 | ~ spl24_107 | spl24_200), inference(avatar_contradiction_clause, [], [f4574])).
fof(f4574, plain, ($false | (~ spl24_81 | ~ spl24_107 | spl24_200)), inference(subsumption_resolution, [], [f4565, f1031])).
fof(f4565, plain, (~ (e0 = op(e1, e3)) | (~ spl24_107 | spl24_200)), inference(backward_demodulation, [], [f1627, f1140])).
fof(f4546, plain, (spl24_109 | ~ spl24_124 | ~ spl24_248), inference(avatar_contradiction_clause, [], [f4545])).
fof(f4545, plain, ($false | (spl24_109 | ~ spl24_124 | ~ spl24_248)), inference(subsumption_resolution, [], [f4541, f1147])).
fof(f1147, plain, (~ (e3 = op(e0, e3)) | spl24_109), inference(avatar_component_clause, [], [f1146])).
fof(f4541, plain, ((e3 = op(e0, e3)) | (~ spl24_124 | ~ spl24_248)), inference(backward_demodulation, [], [f1867, f1211])).
fof(f4418, plain, (spl24_57 | ~ spl24_37 | ~ spl24_155), inference(avatar_split_clause, [], [f4417, f1405, f844, f928])).
fof(f4417, plain, ((e1 = op(e2, e3)) | (~ spl24_37 | ~ spl24_155)), inference(forward_demodulation, [], [f1406, f846])).
fof(f1406, plain, ((op(e2, e3) = op(e3, e2)) | ~ spl24_155), inference(avatar_component_clause, [], [f1405])).
fof(f4366, plain, (~ spl24_41 | ~ spl24_45), inference(avatar_contradiction_clause, [], [f4365])).
fof(f4365, plain, ($false | (~ spl24_41 | ~ spl24_45)), inference(subsumption_resolution, [], [f4364, f485])).
fof(f4364, plain, ((e0 = e4) | (~ spl24_41 | ~ spl24_45)), inference(forward_demodulation, [], [f879, f863])).
fof(f4347, plain, (~ spl24_61 | ~ spl24_64), inference(avatar_contradiction_clause, [], [f4346])).
fof(f4346, plain, ($false | (~ spl24_61 | ~ spl24_64)), inference(subsumption_resolution, [], [f4345, f484])).
fof(f4345, plain, ((e0 = e3) | (~ spl24_61 | ~ spl24_64)), inference(forward_demodulation, [], [f959, f947])).
fof(f4342, plain, (~ spl24_93 | ~ spl24_94), inference(avatar_contradiction_clause, [], [f4341])).
fof(f4341, plain, ($false | (~ spl24_93 | ~ spl24_94)), inference(subsumption_resolution, [], [f4340, f489])).
fof(f4340, plain, ((e2 = e3) | (~ spl24_93 | ~ spl24_94)), inference(forward_demodulation, [], [f1085, f1081])).
fof(f4336, plain, (~ spl24_41 | ~ spl24_107 | spl24_157), inference(avatar_contradiction_clause, [], [f4335])).
fof(f4335, plain, ($false | (~ spl24_41 | ~ spl24_107 | spl24_157)), inference(subsumption_resolution, [], [f4334, f1140])).
fof(f4334, plain, (~ (e1 = op(e0, e3)) | (~ spl24_41 | spl24_157)), inference(forward_demodulation, [], [f1416, f863])).
fof(f4333, plain, (~ spl24_41 | spl24_119 | ~ spl24_158), inference(avatar_contradiction_clause, [], [f4332])).
fof(f4332, plain, ($false | (~ spl24_41 | spl24_119 | ~ spl24_158)), inference(subsumption_resolution, [], [f4331, f1189])).
fof(f1189, plain, (~ (e3 = op(e0, e1)) | spl24_119), inference(avatar_component_clause, [], [f1188])).
fof(f4331, plain, ((e3 = op(e0, e1)) | (~ spl24_41 | ~ spl24_158)), inference(forward_demodulation, [], [f1421, f863])).
fof(f4330, plain, (~ spl24_41 | ~ spl24_81 | spl24_159), inference(avatar_contradiction_clause, [], [f4329])).
fof(f4329, plain, ($false | (~ spl24_41 | ~ spl24_81 | spl24_159)), inference(subsumption_resolution, [], [f4328, f1031])).
fof(f4328, plain, (~ (e0 = op(e1, e3)) | (~ spl24_41 | spl24_159)), inference(forward_demodulation, [], [f1426, f863])).
fof(f4295, plain, (~ spl24_68 | ~ spl24_53), inference(avatar_split_clause, [], [f4294, f911, f974])).
fof(f4294, plain, (~ (e2 = op(e2, e1)) | ~ spl24_53), inference(forward_demodulation, [], [f458, f913])).
fof(f4199, plain, (~ spl24_17 | ~ spl24_77 | spl24_142), inference(avatar_contradiction_clause, [], [f4198])).
fof(f4198, plain, ($false | (~ spl24_17 | ~ spl24_77 | spl24_142)), inference(subsumption_resolution, [], [f4197, f1014])).
fof(f1014, plain, ((e1 = op(e1, e4)) | ~ spl24_77), inference(avatar_component_clause, [], [f1012])).
fof(f4197, plain, (~ (e1 = op(e1, e4)) | (~ spl24_17 | spl24_142)), inference(forward_demodulation, [], [f1343, f762])).
fof(f4150, plain, (~ spl24_48 | ~ spl24_72 | ~ spl24_162), inference(avatar_contradiction_clause, [], [f4149])).
fof(f4149, plain, ($false | (~ spl24_48 | ~ spl24_72 | ~ spl24_162)), inference(subsumption_resolution, [], [f4148, f487])).
fof(f4148, plain, ((e1 = e3) | (~ spl24_48 | ~ spl24_72 | ~ spl24_162)), inference(forward_demodulation, [], [f4146, f993])).
fof(f4146, plain, ((e3 = op(e2, e0)) | (~ spl24_48 | ~ spl24_162)), inference(backward_demodulation, [], [f1440, f892])).
fof(f4137, plain, (~ spl24_72 | ~ spl24_74), inference(avatar_contradiction_clause, [], [f4136])).
fof(f4136, plain, ($false | (~ spl24_72 | ~ spl24_74)), inference(subsumption_resolution, [], [f4135, f487])).
fof(f4135, plain, ((e1 = e3) | (~ spl24_72 | ~ spl24_74)), inference(backward_demodulation, [], [f1001, f993])).
fof(f4133, plain, (~ spl24_74 | ~ spl24_75), inference(avatar_contradiction_clause, [], [f4132])).
fof(f4132, plain, ($false | (~ spl24_74 | ~ spl24_75)), inference(subsumption_resolution, [], [f4131, f491])).
fof(f4131, plain, ((e3 = e4) | (~ spl24_74 | ~ spl24_75)), inference(backward_demodulation, [], [f1005, f1001])).
fof(f4105, plain, (~ spl24_98 | ~ spl24_107 | ~ spl24_115 | ~ spl24_124 | spl24_269), inference(avatar_contradiction_clause, [], [f4104])).
fof(f4104, plain, ($false | (~ spl24_98 | ~ spl24_107 | ~ spl24_115 | ~ spl24_124 | spl24_269)), inference(subsumption_resolution, [], [f4101, f1173])).
fof(f4101, plain, (~ (e4 = op(e0, e2)) | (~ spl24_98 | ~ spl24_107 | ~ spl24_124 | spl24_269)), inference(backward_demodulation, [], [f4090, f1102])).
fof(f4090, plain, (~ (e4 = op(e0, op(e1, e0))) | (~ spl24_107 | ~ spl24_124 | spl24_269)), inference(backward_demodulation, [], [f4066, f1140])).
fof(f4066, plain, (~ (e4 = op(e0, op(op(e0, e3), e0))) | (~ spl24_124 | spl24_269)), inference(backward_demodulation, [], [f2022, f1211])).
fof(f4102, plain, (~ spl24_48 | ~ spl24_98), inference(avatar_split_clause, [], [f4093, f1100, f890])).
fof(f4093, plain, (~ (e2 = op(e3, e0)) | ~ spl24_98), inference(backward_demodulation, [], [f387, f1102])).
fof(f4054, plain, (spl24_89 | ~ spl24_37 | ~ spl24_154), inference(avatar_split_clause, [], [f4053, f1400, f844, f1062])).
fof(f4053, plain, ((e3 = op(e1, e2)) | (~ spl24_37 | ~ spl24_154)), inference(forward_demodulation, [], [f1402, f846])).
fof(f4035, plain, (~ spl24_33 | ~ spl24_37 | spl24_233), inference(avatar_contradiction_clause, [], [f4034])).
fof(f4034, plain, ($false | (~ spl24_33 | ~ spl24_37 | spl24_233)), inference(subsumption_resolution, [], [f4032, f846])).
fof(f4032, plain, (~ (e1 = op(e3, e2)) | (~ spl24_33 | spl24_233)), inference(backward_demodulation, [], [f1785, f829])).
fof(f4025, plain, (~ spl24_21 | ~ spl24_50 | ~ spl24_162), inference(avatar_contradiction_clause, [], [f4024])).
fof(f4024, plain, ($false | (~ spl24_21 | ~ spl24_50 | ~ spl24_162)), inference(subsumption_resolution, [], [f4023, f484])).
fof(f4023, plain, ((e0 = e3) | (~ spl24_21 | ~ spl24_50 | ~ spl24_162)), inference(forward_demodulation, [], [f4017, f779])).
fof(f4017, plain, ((e3 = op(e4, e0)) | (~ spl24_50 | ~ spl24_162)), inference(backward_demodulation, [], [f1440, f900])).
fof(f4001, plain, (~ spl24_108 | ~ spl24_110), inference(avatar_contradiction_clause, [], [f4000])).
fof(f4000, plain, ($false | (~ spl24_108 | ~ spl24_110)), inference(subsumption_resolution, [], [f3999, f490])).
fof(f3999, plain, ((e2 = e4) | (~ spl24_108 | ~ spl24_110)), inference(forward_demodulation, [], [f1152, f1144])).
fof(f3931, plain, (~ spl24_122 | ~ spl24_124), inference(avatar_contradiction_clause, [], [f3930])).
fof(f3930, plain, ($false | (~ spl24_122 | ~ spl24_124)), inference(subsumption_resolution, [], [f3929, f487])).
fof(f3929, plain, ((e1 = e3) | (~ spl24_122 | ~ spl24_124)), inference(forward_demodulation, [], [f1211, f1203])).
fof(f3926, plain, (~ spl24_48 | spl24_74 | ~ spl24_162), inference(avatar_contradiction_clause, [], [f3925])).
fof(f3925, plain, ($false | (~ spl24_48 | spl24_74 | ~ spl24_162)), inference(subsumption_resolution, [], [f3924, f1000])).
fof(f1000, plain, (~ (e3 = op(e2, e0)) | spl24_74), inference(avatar_component_clause, [], [f999])).
fof(f3924, plain, ((e3 = op(e2, e0)) | (~ spl24_48 | ~ spl24_162)), inference(forward_demodulation, [], [f1440, f892])).
fof(f3856, plain, (~ spl24_93 | ~ spl24_95), inference(avatar_contradiction_clause, [], [f3855])).
fof(f3855, plain, ($false | (~ spl24_93 | ~ spl24_95)), inference(subsumption_resolution, [], [f3854, f490])).
fof(f3854, plain, ((e2 = e4) | (~ spl24_93 | ~ spl24_95)), inference(backward_demodulation, [], [f1089, f1081])).
fof(f3835, plain, (~ spl24_114 | ~ spl24_115), inference(avatar_contradiction_clause, [], [f3834])).
fof(f3834, plain, ($false | (~ spl24_114 | ~ spl24_115)), inference(subsumption_resolution, [], [f3833, f491])).
fof(f3833, plain, ((e3 = e4) | (~ spl24_114 | ~ spl24_115)), inference(forward_demodulation, [], [f1173, f1169])).
fof(f3758, plain, (~ spl24_73 | ~ spl24_48), inference(avatar_split_clause, [], [f3757, f890, f995])).
fof(f3757, plain, (~ (e2 = op(e2, e0)) | ~ spl24_48), inference(forward_demodulation, [], [f389, f892])).
fof(f3708, plain, (~ spl24_48 | ~ spl24_50), inference(avatar_contradiction_clause, [], [f3707])).
fof(f3707, plain, ($false | (~ spl24_48 | ~ spl24_50)), inference(subsumption_resolution, [], [f3706, f490])).
fof(f3706, plain, ((e2 = e4) | (~ spl24_48 | ~ spl24_50)), inference(backward_demodulation, [], [f900, f892])).
fof(f3697, plain, (~ spl24_61 | ~ spl24_65), inference(avatar_contradiction_clause, [], [f3696])).
fof(f3696, plain, ($false | (~ spl24_61 | ~ spl24_65)), inference(subsumption_resolution, [], [f3695, f485])).
fof(f3695, plain, ((e0 = e4) | (~ spl24_61 | ~ spl24_65)), inference(backward_demodulation, [], [f963, f947])).
fof(f3672, plain, (~ spl24_99 | ~ spl24_100), inference(avatar_contradiction_clause, [], [f3671])).
fof(f3671, plain, ($false | (~ spl24_99 | ~ spl24_100)), inference(subsumption_resolution, [], [f3670, f491])).
fof(f3670, plain, ((e3 = e4) | (~ spl24_99 | ~ spl24_100)), inference(forward_demodulation, [], [f1110, f1106])).
fof(f3669, plain, (~ spl24_101 | ~ spl24_105), inference(avatar_contradiction_clause, [], [f3668])).
fof(f3668, plain, ($false | (~ spl24_101 | ~ spl24_105)), inference(subsumption_resolution, [], [f3665, f485])).
fof(f3665, plain, ((e0 = e4) | (~ spl24_101 | ~ spl24_105)), inference(backward_demodulation, [], [f1131, f1115])).
fof(f3618, plain, (~ spl24_101 | ~ spl24_21 | spl24_144), inference(avatar_split_clause, [], [f3617, f1350, f777, f1113])).
fof(f3617, plain, (~ (e0 = op(e0, e4)) | (~ spl24_21 | spl24_144)), inference(forward_demodulation, [], [f1352, f779])).
fof(f3613, plain, (~ spl24_101 | ~ spl24_21 | spl24_146), inference(avatar_split_clause, [], [f3612, f1360, f777, f1113])).
fof(f3612, plain, (~ (e0 = op(e0, e4)) | (~ spl24_21 | spl24_146)), inference(forward_demodulation, [], [f1362, f779])).
fof(f3558, plain, (~ spl24_78 | ~ spl24_88), inference(avatar_split_clause, [], [f3557, f1058, f1016])).
fof(f3557, plain, (~ (e2 = op(e1, e4)) | ~ spl24_88), inference(forward_demodulation, [], [f450, f1060])).
fof(f3544, plain, (~ spl24_52 | ~ spl24_57), inference(avatar_split_clause, [], [f3543, f928, f907])).
fof(f3543, plain, (~ (e1 = op(e2, e4)) | ~ spl24_57), inference(forward_demodulation, [], [f461, f930])).
fof(f3534, plain, (~ spl24_52 | ~ spl24_65 | spl24_237), inference(avatar_split_clause, [], [f3402, f1805, f961, f907])).
fof(f3402, plain, (~ (e1 = op(e2, e4)) | (~ spl24_65 | spl24_237)), inference(forward_demodulation, [], [f1807, f963])).
fof(f1807, plain, (~ (e1 = op(e2, op(e2, e2))) | spl24_237), inference(avatar_component_clause, [], [f1805])).
fof(f3508, plain, (~ spl24_33 | ~ spl24_37 | ~ spl24_81 | spl24_241), inference(avatar_contradiction_clause, [], [f3507])).
fof(f3507, plain, ($false | (~ spl24_33 | ~ spl24_37 | ~ spl24_81 | spl24_241)), inference(subsumption_resolution, [], [f3506, f1031])).
fof(f3506, plain, (~ (e0 = op(e1, e3)) | (~ spl24_33 | ~ spl24_37 | spl24_241)), inference(forward_demodulation, [], [f3499, f846])).
fof(f3499, plain, (~ (e0 = op(op(e3, e2), e3)) | (~ spl24_33 | spl24_241)), inference(backward_demodulation, [], [f1836, f829])).
fof(f3487, plain, (~ spl24_47 | ~ spl24_50), inference(avatar_contradiction_clause, [], [f3486])).
fof(f3486, plain, ($false | (~ spl24_47 | ~ spl24_50)), inference(subsumption_resolution, [], [f3485, f488])).
fof(f3485, plain, ((e1 = e4) | (~ spl24_47 | ~ spl24_50)), inference(forward_demodulation, [], [f900, f888])).
fof(f3475, plain, (~ spl24_107 | ~ spl24_109), inference(avatar_contradiction_clause, [], [f3474])).
fof(f3474, plain, ($false | (~ spl24_107 | ~ spl24_109)), inference(subsumption_resolution, [], [f3473, f487])).
fof(f3473, plain, ((e1 = e3) | (~ spl24_107 | ~ spl24_109)), inference(backward_demodulation, [], [f1148, f1140])).
fof(f3446, plain, (~ spl24_29 | ~ spl24_129), inference(avatar_contradiction_clause, [], [f3445])).
fof(f3445, plain, ($false | (~ spl24_29 | ~ spl24_129)), inference(subsumption_resolution, [], [f3444, f491])).
fof(f3444, plain, ((e3 = e4) | (~ spl24_29 | ~ spl24_129)), inference(forward_demodulation, [], [f3423, f812])).
fof(f3423, plain, ((e4 = op(e3, e4)) | ~ spl24_129), inference(backward_demodulation, [], [f329, f1232])).
fof(f3439, plain, (~ spl24_57 | ~ spl24_129), inference(avatar_contradiction_clause, [], [f3438])).
fof(f3438, plain, ($false | (~ spl24_57 | ~ spl24_129)), inference(subsumption_resolution, [], [f3437, f486])).
fof(f3437, plain, ((e1 = e2) | (~ spl24_57 | ~ spl24_129)), inference(forward_demodulation, [], [f3420, f930])).
fof(f3420, plain, ((e2 = op(e2, e3)) | ~ spl24_129), inference(backward_demodulation, [], [f326, f1232])).
fof(f3435, plain, (~ spl24_81 | ~ spl24_129), inference(avatar_contradiction_clause, [], [f3434])).
fof(f3434, plain, ($false | (~ spl24_81 | ~ spl24_129)), inference(subsumption_resolution, [], [f3433, f482])).
fof(f3433, plain, ((e0 = e1) | (~ spl24_81 | ~ spl24_129)), inference(forward_demodulation, [], [f3418, f1031])).
fof(f3418, plain, ((e1 = op(e1, e3)) | ~ spl24_129), inference(backward_demodulation, [], [f324, f1232])).
fof(f3432, plain, (~ spl24_41 | ~ spl24_129), inference(avatar_contradiction_clause, [], [f3431])).
fof(f3431, plain, ($false | (~ spl24_41 | ~ spl24_129)), inference(subsumption_resolution, [], [f3430, f482])).
fof(f3430, plain, ((e0 = e1) | (~ spl24_41 | ~ spl24_129)), inference(forward_demodulation, [], [f3417, f863])).
fof(f3417, plain, ((e1 = op(e3, e1)) | ~ spl24_129), inference(backward_demodulation, [], [f323, f1232])).
fof(f3409, plain, (~ spl24_69 | ~ spl24_78 | ~ spl24_95 | spl24_228), inference(avatar_contradiction_clause, [], [f3408])).
fof(f3408, plain, ($false | (~ spl24_69 | ~ spl24_78 | ~ spl24_95 | spl24_228)), inference(subsumption_resolution, [], [f3407, f980])).
fof(f3407, plain, (~ (e3 = op(e2, e1)) | (~ spl24_78 | ~ spl24_95 | spl24_228)), inference(forward_demodulation, [], [f3406, f1018])).
fof(f3406, plain, (~ (e3 = op(op(e1, e4), e1)) | (~ spl24_95 | spl24_228)), inference(forward_demodulation, [], [f1762, f1089])).
fof(f3397, plain, (~ spl24_119 | ~ spl24_69), inference(avatar_split_clause, [], [f3396, f978, f1188])).
fof(f3396, plain, (~ (e3 = op(e0, e1)) | ~ spl24_69), inference(forward_demodulation, [], [f393, f980])).
fof(f3330, plain, (~ spl24_69 | ~ spl24_78 | ~ spl24_81 | ~ spl24_95 | spl24_227), inference(avatar_contradiction_clause, [], [f3329])).
fof(f3329, plain, ($false | (~ spl24_69 | ~ spl24_78 | ~ spl24_81 | ~ spl24_95 | spl24_227)), inference(subsumption_resolution, [], [f3325, f1031])).
fof(f3325, plain, (~ (e0 = op(e1, e3)) | (~ spl24_69 | ~ spl24_78 | ~ spl24_95 | spl24_227)), inference(backward_demodulation, [], [f3310, f980])).
fof(f3310, plain, (~ (e0 = op(e1, op(e2, e1))) | (~ spl24_78 | ~ spl24_95 | spl24_227)), inference(forward_demodulation, [], [f3304, f1018])).
fof(f3304, plain, (~ (e0 = op(e1, op(op(e1, e4), e1))) | (~ spl24_95 | spl24_227)), inference(backward_demodulation, [], [f1758, f1089])).
fof(f3301, plain, (~ spl24_108 | ~ spl24_109), inference(avatar_contradiction_clause, [], [f3300])).
fof(f3300, plain, ($false | (~ spl24_108 | ~ spl24_109)), inference(subsumption_resolution, [], [f3299, f489])).
fof(f3299, plain, ((e2 = e3) | (~ spl24_108 | ~ spl24_109)), inference(backward_demodulation, [], [f1148, f1144])).
fof(f3282, plain, (~ spl24_51 | ~ spl24_130), inference(avatar_contradiction_clause, [], [f3281])).
fof(f3281, plain, ($false | (~ spl24_51 | ~ spl24_130)), inference(subsumption_resolution, [], [f3280, f483])).
fof(f3280, plain, ((e0 = e2) | (~ spl24_51 | ~ spl24_130)), inference(forward_demodulation, [], [f3267, f905])).
fof(f3267, plain, ((e2 = op(e2, e4)) | ~ spl24_130), inference(backward_demodulation, [], [f326, f1236])).
fof(f3221, plain, (~ spl24_7 | ~ spl24_81 | ~ spl24_133), inference(avatar_contradiction_clause, [], [f3220])).
fof(f3220, plain, ($false | (~ spl24_7 | ~ spl24_81 | ~ spl24_133)), inference(subsumption_resolution, [], [f3219, f485])).
fof(f3219, plain, ((e0 = e4) | (~ spl24_7 | ~ spl24_81 | ~ spl24_133)), inference(forward_demodulation, [], [f3217, f1031])).
fof(f3217, plain, ((e4 = op(e1, e3)) | (~ spl24_7 | ~ spl24_133)), inference(backward_demodulation, [], [f1300, f720])).
fof(f3212, plain, (~ spl24_21 | ~ spl24_25), inference(avatar_contradiction_clause, [], [f3211])).
fof(f3211, plain, ($false | (~ spl24_21 | ~ spl24_25)), inference(subsumption_resolution, [], [f3208, f485])).
fof(f3208, plain, ((e0 = e4) | (~ spl24_21 | ~ spl24_25)), inference(backward_demodulation, [], [f795, f779])).
fof(f3201, plain, (~ spl24_14 | ~ spl24_62 | ~ spl24_70 | spl24_217), inference(avatar_contradiction_clause, [], [f3200])).
fof(f3200, plain, ($false | (~ spl24_14 | ~ spl24_62 | ~ spl24_70 | spl24_217)), inference(subsumption_resolution, [], [f3196, f749])).
fof(f3196, plain, (~ (e3 = op(e4, e2)) | (~ spl24_62 | ~ spl24_70 | spl24_217)), inference(backward_demodulation, [], [f3012, f984])).
fof(f3012, plain, (~ (e3 = op(op(e2, e1), e2)) | (~ spl24_62 | spl24_217)), inference(backward_demodulation, [], [f1712, f951])).
fof(f3190, plain, (~ spl24_81 | ~ spl24_82), inference(avatar_contradiction_clause, [], [f3189])).
fof(f3189, plain, ($false | (~ spl24_81 | ~ spl24_82)), inference(subsumption_resolution, [], [f3188, f482])).
fof(f3188, plain, ((e0 = e1) | (~ spl24_81 | ~ spl24_82)), inference(backward_demodulation, [], [f1035, f1031])).
fof(f3179, plain, (~ spl24_7 | ~ spl24_82), inference(avatar_split_clause, [], [f3170, f1033, f718])).
fof(f3170, plain, (~ (e1 = op(e4, e3)) | ~ spl24_82), inference(backward_demodulation, [], [f418, f1035])).
fof(f3136, plain, (~ spl24_14 | ~ spl24_128), inference(avatar_contradiction_clause, [], [f3135])).
fof(f3135, plain, ($false | (~ spl24_14 | ~ spl24_128)), inference(subsumption_resolution, [], [f3134, f491])).
fof(f3134, plain, ((e3 = e4) | (~ spl24_14 | ~ spl24_128)), inference(forward_demodulation, [], [f3113, f749])).
fof(f3113, plain, ((e4 = op(e4, e2)) | ~ spl24_128), inference(backward_demodulation, [], [f330, f1228])).
fof(f3133, plain, (~ spl24_51 | ~ spl24_128), inference(avatar_contradiction_clause, [], [f3132])).
fof(f3132, plain, ($false | (~ spl24_51 | ~ spl24_128)), inference(subsumption_resolution, [], [f3131, f485])).
fof(f3131, plain, ((e0 = e4) | (~ spl24_51 | ~ spl24_128)), inference(forward_demodulation, [], [f3112, f905])).
fof(f3112, plain, ((e4 = op(e2, e4)) | ~ spl24_128), inference(backward_demodulation, [], [f329, f1228])).
fof(f3062, plain, (~ spl24_7 | ~ spl24_33 | ~ spl24_40 | spl24_225), inference(avatar_split_clause, [], [f3043, f1746, f856, f827, f718])).
fof(f3043, plain, (~ (e1 = op(e4, e3)) | (~ spl24_33 | ~ spl24_40 | spl24_225)), inference(backward_demodulation, [], [f2804, f858])).
fof(f2804, plain, (~ (e1 = op(op(e3, e2), e3)) | (~ spl24_33 | spl24_225)), inference(forward_demodulation, [], [f1748, f829])).
fof(f3061, plain, (~ spl24_6 | ~ spl24_109 | ~ spl24_133), inference(avatar_contradiction_clause, [], [f3060])).
fof(f3060, plain, ($false | (~ spl24_6 | ~ spl24_109 | ~ spl24_133)), inference(subsumption_resolution, [], [f3059, f491])).
fof(f3059, plain, ((e3 = e4) | (~ spl24_6 | ~ spl24_109 | ~ spl24_133)), inference(forward_demodulation, [], [f3058, f1148])).
fof(f3058, plain, ((e4 = op(e0, e3)) | (~ spl24_6 | ~ spl24_133)), inference(backward_demodulation, [], [f1300, f716])).
fof(f3052, plain, (~ spl24_27 | ~ spl24_29), inference(avatar_contradiction_clause, [], [f3051])).
fof(f3051, plain, ($false | (~ spl24_27 | ~ spl24_29)), inference(subsumption_resolution, [], [f3050, f487])).
fof(f3050, plain, ((e1 = e3) | (~ spl24_27 | ~ spl24_29)), inference(backward_demodulation, [], [f812, f804])).
fof(f3045, plain, (~ spl24_6 | ~ spl24_33 | ~ spl24_40 | spl24_241), inference(avatar_split_clause, [], [f3042, f1834, f856, f827, f714])).
fof(f3042, plain, (~ (e0 = op(e4, e3)) | (~ spl24_33 | ~ spl24_40 | spl24_241)), inference(backward_demodulation, [], [f2802, f858])).
fof(f2802, plain, (~ (e0 = op(op(e3, e2), e3)) | (~ spl24_33 | spl24_241)), inference(forward_demodulation, [], [f1836, f829])).
fof(f3040, plain, (~ spl24_36 | ~ spl24_41), inference(avatar_split_clause, [], [f3034, f861, f840])).
fof(f3034, plain, (~ (e0 = op(e3, e2)) | ~ spl24_41), inference(backward_demodulation, [], [f466, f863])).
fof(f3019, plain, (~ spl24_40 | ~ spl24_62 | ~ spl24_69 | spl24_221), inference(avatar_split_clause, [], [f3018, f1728, f978, f949, f856])).
fof(f3018, plain, (~ (e4 = op(e3, e2)) | (~ spl24_62 | ~ spl24_69 | spl24_221)), inference(forward_demodulation, [], [f3013, f980])).
fof(f3013, plain, (~ (e4 = op(op(e2, e1), e2)) | (~ spl24_62 | spl24_221)), inference(backward_demodulation, [], [f1730, f951])).
fof(f2961, plain, (~ spl24_109 | ~ spl24_110), inference(avatar_contradiction_clause, [], [f2960])).
fof(f2960, plain, ($false | (~ spl24_109 | ~ spl24_110)), inference(subsumption_resolution, [], [f2959, f491])).
fof(f2959, plain, ((e3 = e4) | (~ spl24_109 | ~ spl24_110)), inference(forward_demodulation, [], [f1152, f1148])).
fof(f2958, plain, (~ spl24_18 | ~ spl24_127), inference(avatar_contradiction_clause, [], [f2957])).
fof(f2957, plain, ($false | (~ spl24_18 | ~ spl24_127)), inference(subsumption_resolution, [], [f2956, f490])).
fof(f2956, plain, ((e2 = e4) | (~ spl24_18 | ~ spl24_127)), inference(forward_demodulation, [], [f2939, f766])).
fof(f2939, plain, ((e4 = op(e4, e1)) | ~ spl24_127), inference(backward_demodulation, [], [f330, f1224])).
fof(f2913, plain, (~ spl24_16 | ~ spl24_18), inference(avatar_contradiction_clause, [], [f2912])).
fof(f2912, plain, ($false | (~ spl24_16 | ~ spl24_18)), inference(subsumption_resolution, [], [f2911, f483])).
fof(f2911, plain, ((e0 = e2) | (~ spl24_16 | ~ spl24_18)), inference(backward_demodulation, [], [f766, f758])).
fof(f2905, plain, (~ spl24_33 | ~ spl24_35), inference(avatar_contradiction_clause, [], [f2904])).
fof(f2904, plain, ($false | (~ spl24_33 | ~ spl24_35)), inference(subsumption_resolution, [], [f2903, f490])).
fof(f2903, plain, ((e2 = e4) | (~ spl24_33 | ~ spl24_35)), inference(forward_demodulation, [], [f837, f829])).
fof(f2902, plain, (~ spl24_36 | ~ spl24_37), inference(avatar_contradiction_clause, [], [f2901])).
fof(f2901, plain, ($false | (~ spl24_36 | ~ spl24_37)), inference(subsumption_resolution, [], [f2900, f482])).
fof(f2900, plain, ((e0 = e1) | (~ spl24_36 | ~ spl24_37)), inference(backward_demodulation, [], [f846, f842])).
fof(f2899, plain, (~ spl24_28 | ~ spl24_43), inference(avatar_split_clause, [], [f2897, f869, f806])).
fof(f2897, plain, (~ (e2 = op(e3, e4)) | ~ spl24_43), inference(backward_demodulation, [], [f468, f871])).
fof(f2881, plain, (~ spl24_69 | ~ spl24_70), inference(avatar_contradiction_clause, [], [f2880])).
fof(f2880, plain, ($false | (~ spl24_69 | ~ spl24_70)), inference(subsumption_resolution, [], [f2879, f491])).
fof(f2879, plain, ((e3 = e4) | (~ spl24_69 | ~ spl24_70)), inference(forward_demodulation, [], [f984, f980])).
fof(f2878, plain, (~ spl24_76 | ~ spl24_79), inference(avatar_contradiction_clause, [], [f2877])).
fof(f2877, plain, ($false | (~ spl24_76 | ~ spl24_79)), inference(subsumption_resolution, [], [f2876, f484])).
fof(f2876, plain, ((e0 = e3) | (~ spl24_76 | ~ spl24_79)), inference(backward_demodulation, [], [f1022, f1010])).
fof(f2818, plain, (~ spl24_54 | ~ spl24_2 | ~ spl24_18 | spl24_219), inference(avatar_split_clause, [], [f2777, f1719, f764, f697, f915])).
fof(f2777, plain, (~ (e3 = op(e2, e4)) | (~ spl24_2 | ~ spl24_18 | spl24_219)), inference(backward_demodulation, [], [f2392, f766])).
fof(f2392, plain, (~ (e3 = op(op(e4, e1), e4)) | (~ spl24_2 | spl24_219)), inference(backward_demodulation, [], [f1721, f699])).
fof(f2808, plain, (~ spl24_85 | ~ spl24_33 | ~ spl24_37 | spl24_223), inference(avatar_split_clause, [], [f2807, f1737, f844, f827, f1045])).
fof(f2807, plain, (~ (e4 = op(e1, e3)) | (~ spl24_33 | ~ spl24_37 | spl24_223)), inference(forward_demodulation, [], [f2806, f846])).
fof(f2806, plain, (~ (e4 = op(op(e3, e2), e3)) | (~ spl24_33 | spl24_223)), inference(forward_demodulation, [], [f1739, f829])).
fof(f2791, plain, (~ spl24_9 | ~ spl24_14), inference(avatar_split_clause, [], [f2790, f747, f726])).
fof(f2790, plain, (~ (e3 = op(e4, e3)) | ~ spl24_14), inference(backward_demodulation, [], [f479, f749])).
fof(f2786, plain, (~ spl24_11 | ~ spl24_14), inference(avatar_contradiction_clause, [], [f2785])).
fof(f2785, plain, ($false | (~ spl24_11 | ~ spl24_14)), inference(subsumption_resolution, [], [f2783, f484])).
fof(f2783, plain, ((e0 = e3) | (~ spl24_11 | ~ spl24_14)), inference(backward_demodulation, [], [f749, f737])).
fof(f2775, plain, (~ spl24_31 | ~ spl24_32), inference(avatar_contradiction_clause, [], [f2774])).
fof(f2774, plain, ($false | (~ spl24_31 | ~ spl24_32)), inference(subsumption_resolution, [], [f2773, f482])).
fof(f2773, plain, ((e0 = e1) | (~ spl24_31 | ~ spl24_32)), inference(backward_demodulation, [], [f825, f821])).
fof(f2772, plain, (~ spl24_32 | ~ spl24_33), inference(avatar_contradiction_clause, [], [f2771])).
fof(f2771, plain, ($false | (~ spl24_32 | ~ spl24_33)), inference(subsumption_resolution, [], [f2770, f486])).
fof(f2770, plain, ((e1 = e2) | (~ spl24_32 | ~ spl24_33)), inference(forward_demodulation, [], [f829, f825])).
fof(f2764, plain, (~ spl24_51 | ~ spl24_54), inference(avatar_contradiction_clause, [], [f2763])).
fof(f2763, plain, ($false | (~ spl24_51 | ~ spl24_54)), inference(subsumption_resolution, [], [f2761, f484])).
fof(f2761, plain, ((e0 = e3) | (~ spl24_51 | ~ spl24_54)), inference(backward_demodulation, [], [f917, f905])).
fof(f2711, plain, (~ spl24_63 | ~ spl24_73), inference(avatar_split_clause, [], [f2710, f995, f953])).
fof(f2710, plain, (~ (e2 = op(e2, e2)) | ~ spl24_73), inference(forward_demodulation, [], [f453, f997])).
fof(f2686, plain, (~ spl24_44 | ~ spl24_29), inference(avatar_split_clause, [], [f2681, f810, f873])).
fof(f2681, plain, (~ (e3 = op(e3, e1)) | ~ spl24_29), inference(backward_demodulation, [], [f468, f812])).
fof(f2671, plain, (~ spl24_19 | ~ spl24_2 | spl24_215), inference(avatar_split_clause, [], [f2391, f1701, f697, f768])).
fof(f2391, plain, (~ (e3 = op(e4, e1)) | (~ spl24_2 | spl24_215)), inference(backward_demodulation, [], [f1703, f699])).
fof(f2658, plain, (~ spl24_29 | ~ spl24_9 | spl24_134), inference(avatar_split_clause, [], [f2649, f1303, f726, f810])).
fof(f2649, plain, (~ (e3 = op(e3, e4)) | (~ spl24_9 | spl24_134)), inference(backward_demodulation, [], [f1305, f728])).
fof(f728, plain, ((e3 = op(e4, e3)) | ~ spl24_9), inference(avatar_component_clause, [], [f726])).
fof(f2652, plain, (~ spl24_19 | ~ spl24_9), inference(avatar_split_clause, [], [f2644, f726, f768])).
fof(f2644, plain, (~ (e3 = op(e4, e1)) | ~ spl24_9), inference(backward_demodulation, [], [f477, f728])).
fof(f2623, plain, (~ spl24_32 | ~ spl24_43 | ~ spl24_60 | spl24_223), inference(avatar_contradiction_clause, [], [f2622])).
fof(f2622, plain, ($false | (~ spl24_32 | ~ spl24_43 | ~ spl24_60 | spl24_223)), inference(subsumption_resolution, [], [f2621, f942])).
fof(f2621, plain, (~ (e4 = op(e2, e3)) | (~ spl24_32 | ~ spl24_43 | spl24_223)), inference(forward_demodulation, [], [f2616, f871])).
fof(f2616, plain, (~ (e4 = op(op(e3, e1), e3)) | (~ spl24_32 | spl24_223)), inference(backward_demodulation, [], [f1739, f825])).
fof(f2612, plain, (~ spl24_43 | ~ spl24_45), inference(avatar_contradiction_clause, [], [f2611])).
fof(f2611, plain, ($false | (~ spl24_43 | ~ spl24_45)), inference(subsumption_resolution, [], [f2610, f490])).
fof(f2610, plain, ((e2 = e4) | (~ spl24_43 | ~ spl24_45)), inference(backward_demodulation, [], [f879, f871])).
fof(f2605, plain, (~ spl24_34 | ~ spl24_49), inference(avatar_split_clause, [], [f2602, f894, f831])).
fof(f2602, plain, (~ (e3 = op(e3, e3)) | ~ spl24_49), inference(backward_demodulation, [], [f464, f896])).
fof(f2603, plain, (~ spl24_44 | ~ spl24_49), inference(avatar_split_clause, [], [f2600, f894, f873])).
fof(f2600, plain, (~ (e3 = op(e3, e1)) | ~ spl24_49), inference(backward_demodulation, [], [f462, f896])).
fof(f2595, plain, (~ spl24_57 | ~ spl24_67), inference(avatar_split_clause, [], [f2593, f970, f928])).
fof(f2593, plain, (~ (e1 = op(e2, e3)) | ~ spl24_67), inference(backward_demodulation, [], [f457, f972])).
fof(f2588, plain, (~ spl24_58 | ~ spl24_73), inference(avatar_split_clause, [], [f2582, f995, f932])).
fof(f2582, plain, (~ (e2 = op(e2, e3)) | ~ spl24_73), inference(backward_demodulation, [], [f454, f997])).
fof(f2577, plain, (~ spl24_83 | ~ spl24_85), inference(avatar_contradiction_clause, [], [f2576])).
fof(f2576, plain, ($false | (~ spl24_83 | ~ spl24_85)), inference(subsumption_resolution, [], [f2575, f490])).
fof(f2575, plain, ((e2 = e4) | (~ spl24_83 | ~ spl24_85)), inference(backward_demodulation, [], [f1047, f1039])).
fof(f2556, plain, (~ spl24_47 | ~ spl24_97), inference(avatar_split_clause, [], [f2548, f1096, f886])).
fof(f2548, plain, (~ (e1 = op(e3, e0)) | ~ spl24_97), inference(backward_demodulation, [], [f387, f1098])).
fof(f2555, plain, (~ spl24_72 | ~ spl24_97), inference(avatar_split_clause, [], [f2547, f1096, f991])).
fof(f2547, plain, (~ (e1 = op(e2, e0)) | ~ spl24_97), inference(backward_demodulation, [], [f386, f1098])).
fof(f2545, plain, (~ spl24_34 | ~ spl24_109), inference(avatar_split_clause, [], [f2542, f1146, f831])).
fof(f2542, plain, (~ (e3 = op(e3, e3)) | ~ spl24_109), inference(backward_demodulation, [], [f414, f1148])).
fof(f2532, plain, (~ spl24_112 | ~ spl24_117), inference(avatar_split_clause, [], [f2528, f1180, f1159])).
fof(f2528, plain, (~ (e1 = op(e0, e2)) | ~ spl24_117), inference(backward_demodulation, [], [f436, f1182])).
fof(f2531, plain, (~ spl24_42 | ~ spl24_117), inference(avatar_split_clause, [], [f2527, f1180, f865])).
fof(f2527, plain, (~ (e1 = op(e3, e1)) | ~ spl24_117), inference(backward_demodulation, [], [f394, f1182])).
fof(f2509, plain, (~ spl24_103 | ~ spl24_105), inference(avatar_contradiction_clause, [], [f2508])).
fof(f2508, plain, ($false | (~ spl24_103 | ~ spl24_105)), inference(subsumption_resolution, [], [f2507, f490])).
fof(f2507, plain, ((e2 = e4) | (~ spl24_103 | ~ spl24_105)), inference(forward_demodulation, [], [f1131, f1123])).
fof(f2500, plain, (~ spl24_86 | ~ spl24_61), inference(avatar_split_clause, [], [f2499, f945, f1050])).
fof(f2499, plain, (~ (e0 = op(e1, e2)) | ~ spl24_61), inference(forward_demodulation, [], [f406, f947])).
fof(f2494, plain, (~ spl24_81 | ~ spl24_6), inference(avatar_split_clause, [], [f2493, f714, f1029])).
fof(f2493, plain, (~ (e0 = op(e1, e3)) | ~ spl24_6), inference(forward_demodulation, [], [f418, f716])).
fof(f2487, plain, (~ spl24_76 | ~ spl24_26), inference(avatar_split_clause, [], [f2486, f798, f1008])).
fof(f2486, plain, (~ (e0 = op(e1, e4)) | ~ spl24_26), inference(forward_demodulation, [], [f427, f800])).
fof(f2472, plain, (~ spl24_67 | ~ spl24_17), inference(avatar_split_clause, [], [f2471, f760, f970])).
fof(f2471, plain, (~ (e1 = op(e2, e1)) | ~ spl24_17), inference(forward_demodulation, [], [f400, f762])).
fof(f2470, plain, (~ spl24_66 | ~ spl24_61), inference(avatar_split_clause, [], [f2469, f945, f966])).
fof(f2469, plain, (~ (e0 = op(e2, e1)) | ~ spl24_61), inference(forward_demodulation, [], [f456, f947])).
fof(f2468, plain, (~ spl24_67 | ~ spl24_52), inference(avatar_split_clause, [], [f2467, f907, f970])).
fof(f2467, plain, (~ (e1 = op(e2, e1)) | ~ spl24_52), inference(forward_demodulation, [], [f458, f909])).
fof(f2457, plain, (~ spl24_52 | ~ spl24_54), inference(avatar_contradiction_clause, [], [f2456])).
fof(f2456, plain, ($false | (~ spl24_52 | ~ spl24_54)), inference(subsumption_resolution, [], [f2455, f487])).
fof(f2455, plain, ((e1 = e3) | (~ spl24_52 | ~ spl24_54)), inference(forward_demodulation, [], [f917, f909])).
fof(f2452, plain, (~ spl24_46 | ~ spl24_26), inference(avatar_split_clause, [], [f2451, f798, f882])).
fof(f2451, plain, (~ (e0 = op(e3, e0)) | ~ spl24_26), inference(forward_demodulation, [], [f465, f800])).
fof(f2448, plain, (~ spl24_41 | ~ spl24_26), inference(avatar_split_clause, [], [f2447, f798, f861])).
fof(f2447, plain, (~ (e0 = op(e3, e1)) | ~ spl24_26), inference(forward_demodulation, [], [f468, f800])).
fof(f2436, plain, (~ spl24_6 | ~ spl24_26 | spl24_134), inference(avatar_contradiction_clause, [], [f2435])).
fof(f2435, plain, ($false | (~ spl24_6 | ~ spl24_26 | spl24_134)), inference(subsumption_resolution, [], [f2434, f800])).
fof(f2434, plain, (~ (e0 = op(e3, e4)) | (~ spl24_6 | spl24_134)), inference(forward_demodulation, [], [f1305, f716])).
fof(f2430, plain, (~ spl24_23 | ~ spl24_25), inference(avatar_contradiction_clause, [], [f2429])).
fof(f2429, plain, ($false | (~ spl24_23 | ~ spl24_25)), inference(subsumption_resolution, [], [f2428, f490])).
fof(f2428, plain, ((e2 = e4) | (~ spl24_23 | ~ spl24_25)), inference(forward_demodulation, [], [f795, f787])).
fof(f2416, plain, (~ spl24_11 | ~ spl24_6), inference(avatar_split_clause, [], [f2415, f714, f735])).
fof(f2415, plain, (~ (e0 = op(e4, e2)) | ~ spl24_6), inference(forward_demodulation, [], [f479, f716])).
fof(f2411, plain, (~ spl24_12 | ~ spl24_14), inference(avatar_contradiction_clause, [], [f2410])).
fof(f2410, plain, ($false | (~ spl24_12 | ~ spl24_14)), inference(subsumption_resolution, [], [f2409, f487])).
fof(f2409, plain, ((e1 = e3) | (~ spl24_12 | ~ spl24_14)), inference(forward_demodulation, [], [f749, f741])).
fof(f2404, plain, (spl24_110 | ~ spl24_6 | ~ spl24_133), inference(avatar_split_clause, [], [f2403, f1298, f714, f1150])).
fof(f2403, plain, ((e4 = op(e0, e3)) | (~ spl24_6 | ~ spl24_133)), inference(forward_demodulation, [], [f1300, f716])).
fof(f2402, plain, (~ spl24_6 | ~ spl24_9), inference(avatar_contradiction_clause, [], [f2401])).
fof(f2401, plain, ($false | (~ spl24_6 | ~ spl24_9)), inference(subsumption_resolution, [], [f2400, f484])).
fof(f2400, plain, ((e0 = e3) | (~ spl24_6 | ~ spl24_9)), inference(forward_demodulation, [], [f728, f716])).
fof(f2388, plain, (~ spl24_6 | ~ spl24_7), inference(avatar_contradiction_clause, [], [f2387])).
fof(f2387, plain, ($false | (~ spl24_6 | ~ spl24_7)), inference(subsumption_resolution, [], [f2386, f482])).
fof(f2386, plain, ((e0 = e1) | (~ spl24_6 | ~ spl24_7)), inference(backward_demodulation, [], [f720, f716])).
fof(f2385, plain, (~ spl24_7 | ~ spl24_8), inference(avatar_contradiction_clause, [], [f2384])).
fof(f2384, plain, ($false | (~ spl24_7 | ~ spl24_8)), inference(subsumption_resolution, [], [f2383, f486])).
fof(f2383, plain, ((e1 = e2) | (~ spl24_7 | ~ spl24_8)), inference(backward_demodulation, [], [f724, f720])).
fof(f2382, plain, (~ spl24_8 | ~ spl24_9), inference(avatar_contradiction_clause, [], [f2381])).
fof(f2381, plain, ($false | (~ spl24_8 | ~ spl24_9)), inference(subsumption_resolution, [], [f2380, f489])).
fof(f2380, plain, ((e2 = e3) | (~ spl24_8 | ~ spl24_9)), inference(backward_demodulation, [], [f728, f724])).
fof(f2379, plain, (~ spl24_9 | ~ spl24_10), inference(avatar_contradiction_clause, [], [f2378])).
fof(f2378, plain, ($false | (~ spl24_9 | ~ spl24_10)), inference(subsumption_resolution, [], [f2377, f491])).
fof(f2377, plain, ((e3 = e4) | (~ spl24_9 | ~ spl24_10)), inference(backward_demodulation, [], [f732, f728])).
fof(f2375, plain, (~ spl24_5 | ~ spl24_10), inference(avatar_split_clause, [], [f2372, f730, f709])).
fof(f2372, plain, (~ (e4 = op(e4, e4)) | ~ spl24_10), inference(backward_demodulation, [], [f481, f732])).
fof(f2371, plain, (~ spl24_12 | ~ spl24_13), inference(avatar_contradiction_clause, [], [f2370])).
fof(f2370, plain, ($false | (~ spl24_12 | ~ spl24_13)), inference(subsumption_resolution, [], [f2369, f486])).
fof(f2369, plain, ((e1 = e2) | (~ spl24_12 | ~ spl24_13)), inference(backward_demodulation, [], [f745, f741])).
fof(f2362, plain, (~ spl24_5 | ~ spl24_15), inference(avatar_split_clause, [], [f2360, f751, f709])).
fof(f2360, plain, (~ (e4 = op(e4, e4)) | ~ spl24_15), inference(backward_demodulation, [], [f480, f753])).
fof(f2355, plain, (~ spl24_18 | ~ spl24_19), inference(avatar_contradiction_clause, [], [f2354])).
fof(f2354, plain, ($false | (~ spl24_18 | ~ spl24_19)), inference(subsumption_resolution, [], [f2353, f489])).
fof(f2353, plain, ((e2 = e3) | (~ spl24_18 | ~ spl24_19)), inference(backward_demodulation, [], [f770, f766])).
fof(f2352, plain, (~ spl24_19 | ~ spl24_20), inference(avatar_contradiction_clause, [], [f2351])).
fof(f2351, plain, ($false | (~ spl24_19 | ~ spl24_20)), inference(subsumption_resolution, [], [f2350, f491])).
fof(f2350, plain, ((e3 = e4) | (~ spl24_19 | ~ spl24_20)), inference(backward_demodulation, [], [f774, f770])).
fof(f2343, plain, (~ spl24_22 | ~ spl24_23), inference(avatar_contradiction_clause, [], [f2342])).
fof(f2342, plain, ($false | (~ spl24_22 | ~ spl24_23)), inference(subsumption_resolution, [], [f2341, f486])).
fof(f2341, plain, ((e1 = e2) | (~ spl24_22 | ~ spl24_23)), inference(backward_demodulation, [], [f787, f783])).
fof(f2340, plain, (~ spl24_23 | ~ spl24_24), inference(avatar_contradiction_clause, [], [f2339])).
fof(f2339, plain, ($false | (~ spl24_23 | ~ spl24_24)), inference(subsumption_resolution, [], [f2338, f489])).
fof(f2338, plain, ((e2 = e3) | (~ spl24_23 | ~ spl24_24)), inference(backward_demodulation, [], [f791, f787])).
fof(f2337, plain, (~ spl24_24 | ~ spl24_25), inference(avatar_contradiction_clause, [], [f2336])).
fof(f2336, plain, ($false | (~ spl24_24 | ~ spl24_25)), inference(subsumption_resolution, [], [f2335, f491])).
fof(f2335, plain, ((e3 = e4) | (~ spl24_24 | ~ spl24_25)), inference(backward_demodulation, [], [f795, f791])).
fof(f2333, plain, (~ spl24_10 | ~ spl24_25), inference(avatar_split_clause, [], [f2329, f793, f730])).
fof(f2329, plain, (~ (e4 = op(e4, e3)) | ~ spl24_25), inference(backward_demodulation, [], [f474, f795])).
fof(f2326, plain, (~ spl24_26 | ~ spl24_27), inference(avatar_contradiction_clause, [], [f2325])).
fof(f2325, plain, ($false | (~ spl24_26 | ~ spl24_27)), inference(subsumption_resolution, [], [f2324, f482])).
fof(f2324, plain, ((e0 = e1) | (~ spl24_26 | ~ spl24_27)), inference(backward_demodulation, [], [f804, f800])).
fof(f2323, plain, (~ spl24_27 | ~ spl24_28), inference(avatar_contradiction_clause, [], [f2322])).
fof(f2322, plain, ($false | (~ spl24_27 | ~ spl24_28)), inference(subsumption_resolution, [], [f2321, f486])).
fof(f2321, plain, ((e1 = e2) | (~ spl24_27 | ~ spl24_28)), inference(backward_demodulation, [], [f808, f804])).
fof(f2320, plain, (~ spl24_28 | ~ spl24_29), inference(avatar_contradiction_clause, [], [f2319])).
fof(f2319, plain, ($false | (~ spl24_28 | ~ spl24_29)), inference(subsumption_resolution, [], [f2318, f489])).
fof(f2318, plain, ((e2 = e3) | (~ spl24_28 | ~ spl24_29)), inference(backward_demodulation, [], [f812, f808])).
fof(f2314, plain, (~ spl24_10 | ~ spl24_30 | spl24_134), inference(avatar_split_clause, [], [f2312, f1303, f814, f730])).
fof(f2312, plain, (~ (e4 = op(e4, e3)) | (~ spl24_30 | spl24_134)), inference(backward_demodulation, [], [f1305, f816])).
fof(f2292, plain, (~ spl24_31 | ~ spl24_36), inference(avatar_split_clause, [], [f2289, f840, f819])).
fof(f2289, plain, (~ (e0 = op(e3, e3)) | ~ spl24_36), inference(backward_demodulation, [], [f469, f842])).
fof(f2277, plain, (~ spl24_36 | ~ spl24_46), inference(avatar_split_clause, [], [f2272, f882, f840])).
fof(f2272, plain, (~ (e0 = op(e3, e2)) | ~ spl24_46), inference(backward_demodulation, [], [f463, f884])).
fof(f2269, plain, (~ spl24_52 | ~ spl24_53), inference(avatar_contradiction_clause, [], [f2268])).
fof(f2268, plain, ($false | (~ spl24_52 | ~ spl24_53)), inference(subsumption_resolution, [], [f2267, f486])).
fof(f2267, plain, ((e1 = e2) | (~ spl24_52 | ~ spl24_53)), inference(backward_demodulation, [], [f913, f909])).
fof(f2263, plain, (~ spl24_54 | ~ spl24_55), inference(avatar_contradiction_clause, [], [f2262])).
fof(f2262, plain, ($false | (~ spl24_54 | ~ spl24_55)), inference(subsumption_resolution, [], [f2261, f491])).
fof(f2261, plain, ((e3 = e4) | (~ spl24_54 | ~ spl24_55)), inference(backward_demodulation, [], [f921, f917])).
fof(f2225, plain, (~ spl24_41 | ~ spl24_66), inference(avatar_split_clause, [], [f2220, f966, f861])).
fof(f2220, plain, (~ (e0 = op(e3, e1)) | ~ spl24_66), inference(backward_demodulation, [], [f399, f968])).
fof(f2218, plain, (~ spl24_56 | ~ spl24_71), inference(avatar_split_clause, [], [f2212, f987, f924])).
fof(f2212, plain, (~ (e0 = op(e2, e3)) | ~ spl24_71), inference(backward_demodulation, [], [f454, f989])).
fof(f2216, plain, (~ spl24_66 | ~ spl24_71), inference(avatar_split_clause, [], [f2210, f987, f966])).
fof(f2210, plain, (~ (e0 = op(e2, e1)) | ~ spl24_71), inference(backward_demodulation, [], [f452, f989])).
fof(f2215, plain, (~ spl24_21 | ~ spl24_71), inference(avatar_split_clause, [], [f2209, f987, f777])).
fof(f2209, plain, (~ (e0 = op(e4, e0)) | ~ spl24_71), inference(backward_demodulation, [], [f390, f989])).
fof(f2214, plain, (~ spl24_46 | ~ spl24_71), inference(avatar_split_clause, [], [f2208, f987, f882])).
fof(f2208, plain, (~ (e0 = op(e3, e0)) | ~ spl24_71), inference(backward_demodulation, [], [f389, f989])).
fof(f2204, plain, (~ spl24_78 | ~ spl24_79), inference(avatar_contradiction_clause, [], [f2203])).
fof(f2203, plain, ($false | (~ spl24_78 | ~ spl24_79)), inference(subsumption_resolution, [], [f2202, f489])).
fof(f2202, plain, ((e2 = e3) | (~ spl24_78 | ~ spl24_79)), inference(backward_demodulation, [], [f1022, f1018])).
fof(f2201, plain, (~ spl24_79 | ~ spl24_80), inference(avatar_contradiction_clause, [], [f2200])).
fof(f2200, plain, ($false | (~ spl24_79 | ~ spl24_80)), inference(subsumption_resolution, [], [f2199, f491])).
fof(f2199, plain, ((e3 = e4) | (~ spl24_79 | ~ spl24_80)), inference(backward_demodulation, [], [f1026, f1022])).
fof(f2192, plain, (~ spl24_76 | ~ spl24_81), inference(avatar_split_clause, [], [f2188, f1029, f1008])).
fof(f2188, plain, (~ (e0 = op(e1, e4)) | ~ spl24_81), inference(backward_demodulation, [], [f451, f1031])).
fof(f2190, plain, (~ spl24_31 | ~ spl24_81), inference(avatar_split_clause, [], [f2186, f1029, f819])).
fof(f2186, plain, (~ (e0 = op(e3, e3)) | ~ spl24_81), inference(backward_demodulation, [], [f417, f1031])).
fof(f2189, plain, (~ spl24_56 | ~ spl24_81), inference(avatar_split_clause, [], [f2185, f1029, f924])).
fof(f2185, plain, (~ (e0 = op(e2, e3)) | ~ spl24_81), inference(backward_demodulation, [], [f416, f1031])).
fof(f2183, plain, (~ spl24_81 | ~ spl24_86), inference(avatar_split_clause, [], [f2178, f1050, f1029])).
fof(f2178, plain, (~ (e0 = op(e1, e3)) | ~ spl24_86), inference(backward_demodulation, [], [f449, f1052])).
fof(f2160, plain, (~ spl24_81 | ~ spl24_91), inference(avatar_split_clause, [], [f2150, f1071, f1029])).
fof(f2150, plain, (~ (e0 = op(e1, e3)) | ~ spl24_91), inference(backward_demodulation, [], [f447, f1073])).
fof(f2145, plain, (~ spl24_76 | ~ spl24_96), inference(avatar_split_clause, [], [f2138, f1092, f1008])).
fof(f2138, plain, (~ (e0 = op(e1, e4)) | ~ spl24_96), inference(backward_demodulation, [], [f445, f1094])).
fof(f2144, plain, (~ spl24_81 | ~ spl24_96), inference(avatar_split_clause, [], [f2137, f1092, f1029])).
fof(f2137, plain, (~ (e0 = op(e1, e3)) | ~ spl24_96), inference(backward_demodulation, [], [f444, f1094])).
fof(f2143, plain, (~ spl24_86 | ~ spl24_96), inference(avatar_split_clause, [], [f2136, f1092, f1050])).
fof(f2136, plain, (~ (e0 = op(e1, e2)) | ~ spl24_96), inference(backward_demodulation, [], [f443, f1094])).
fof(f2141, plain, (~ spl24_21 | ~ spl24_96), inference(avatar_split_clause, [], [f2134, f1092, f777])).
fof(f2134, plain, (~ (e0 = op(e4, e0)) | ~ spl24_96), inference(backward_demodulation, [], [f388, f1094])).
fof(f2140, plain, (~ spl24_46 | ~ spl24_96), inference(avatar_split_clause, [], [f2133, f1092, f882])).
fof(f2133, plain, (~ (e0 = op(e3, e0)) | ~ spl24_96), inference(backward_demodulation, [], [f387, f1094])).
fof(f2139, plain, (~ spl24_71 | ~ spl24_96), inference(avatar_split_clause, [], [f2132, f1092, f987])).
fof(f2132, plain, (~ (e0 = op(e2, e0)) | ~ spl24_96), inference(backward_demodulation, [], [f386, f1094])).
fof(f2131, plain, (~ spl24_102 | ~ spl24_103), inference(avatar_contradiction_clause, [], [f2130])).
fof(f2130, plain, ($false | (~ spl24_102 | ~ spl24_103)), inference(subsumption_resolution, [], [f2129, f486])).
fof(f2129, plain, ((e1 = e2) | (~ spl24_102 | ~ spl24_103)), inference(backward_demodulation, [], [f1123, f1119])).
fof(f2128, plain, (~ spl24_103 | ~ spl24_104), inference(avatar_contradiction_clause, [], [f2127])).
fof(f2127, plain, ($false | (~ spl24_103 | ~ spl24_104)), inference(subsumption_resolution, [], [f2126, f489])).
fof(f2126, plain, ((e2 = e3) | (~ spl24_103 | ~ spl24_104)), inference(backward_demodulation, [], [f1127, f1123])).
fof(f2125, plain, (~ spl24_104 | ~ spl24_105), inference(avatar_contradiction_clause, [], [f2124])).
fof(f2124, plain, ($false | (~ spl24_104 | ~ spl24_105)), inference(subsumption_resolution, [], [f2123, f491])).
fof(f2123, plain, ((e3 = e4) | (~ spl24_104 | ~ spl24_105)), inference(backward_demodulation, [], [f1131, f1127])).
fof(f2121, plain, (~ spl24_30 | ~ spl24_105), inference(avatar_split_clause, [], [f2117, f1129, f814])).
fof(f2117, plain, (~ (e4 = op(e3, e4)) | ~ spl24_105), inference(backward_demodulation, [], [f424, f1131])).
fof(f2111, plain, (~ spl24_56 | ~ spl24_106), inference(avatar_split_clause, [], [f2106, f1134, f924])).
fof(f2106, plain, (~ (e0 = op(e2, e3)) | ~ spl24_106), inference(backward_demodulation, [], [f413, f1136])).
fof(f2110, plain, (~ spl24_81 | ~ spl24_106), inference(avatar_split_clause, [], [f2105, f1134, f1029])).
fof(f2105, plain, (~ (e0 = op(e1, e3)) | ~ spl24_106), inference(backward_demodulation, [], [f412, f1136])).
fof(f2101, plain, (~ spl24_36 | ~ spl24_111), inference(avatar_split_clause, [], [f2095, f1155, f840])).
fof(f2095, plain, (~ (e0 = op(e3, e2)) | ~ spl24_111), inference(backward_demodulation, [], [f404, f1157])).
fof(f2090, plain, (~ spl24_111 | ~ spl24_116), inference(avatar_split_clause, [], [f2083, f1176, f1155])).
fof(f2083, plain, (~ (e0 = op(e0, e2)) | ~ spl24_116), inference(backward_demodulation, [], [f436, f1178])).
fof(f2087, plain, (~ spl24_66 | ~ spl24_116), inference(avatar_split_clause, [], [f2080, f1176, f966])).
fof(f2080, plain, (~ (e0 = op(e2, e1)) | ~ spl24_116), inference(backward_demodulation, [], [f393, f1178])).
fof(f2077, plain, (~ spl24_106 | ~ spl24_121), inference(avatar_split_clause, [], [f2065, f1197, f1134])).
fof(f2065, plain, (~ (e0 = op(e0, e3)) | ~ spl24_121), inference(backward_demodulation, [], [f434, f1199])).
fof(f2073, plain, (~ spl24_46 | ~ spl24_121), inference(avatar_split_clause, [], [f2061, f1197, f882])).
fof(f2061, plain, (~ (e0 = op(e3, e0)) | ~ spl24_121), inference(backward_demodulation, [], [f384, f1199])).
fof(f2038, plain, (~ spl24_268 | ~ spl24_94 | ~ spl24_235 | ~ spl24_261), inference(avatar_split_clause, [], [f686, f1949, f1794, f1083, f2010])).
fof(f686, plain, (~ (e0 = op(e1, op(e1, e1))) | ~ (e2 = op(op(e1, op(e1, e1)), e1)) | ~ (e3 = op(e1, e1)) | ~ (e4 = op(e1, op(op(e1, op(e1, e1)), e1)))), inference(cnf_transformation, [], [f246])).
fof(f246, plain, (~ (e0 = op(e1, op(e1, e1))) | ~ (e2 = op(op(e1, op(e1, e1)), e1)) | ~ (e3 = op(e1, e1)) | ~ (e4 = op(e1, op(op(e1, op(e1, e1)), e1)))), inference(ennf_transformation, [], [f126])).
fof(f126, plain, ~ ((e0 = op(e1, op(e1, e1))) & (e2 = op(op(e1, op(e1, e1)), e1)) & (e3 = op(e1, e1)) & (e4 = op(e1, op(op(e1, op(e1, e1)), e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax126)).
fof(f2037, plain, (~ spl24_269 | ~ spl24_124 | ~ spl24_251 | ~ spl24_260), inference(avatar_split_clause, [], [f685, f1944, f1882, f1209, f2020])).
fof(f685, plain, (~ (e1 = op(e0, op(e0, e0))) | ~ (e2 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e3) | ~ (e4 = op(e0, op(op(e0, op(e0, e0)), e0)))), inference(cnf_transformation, [], [f245])).
fof(f245, plain, (~ (e1 = op(e0, op(e0, e0))) | ~ (e2 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e3) | ~ (e4 = op(e0, op(op(e0, op(e0, e0)), e0)))), inference(ennf_transformation, [], [f125])).
fof(f125, plain, ~ ((e1 = op(e0, op(e0, e0))) & (e2 = op(op(e0, op(e0, e0)), e0)) & (op(e0, e0) = e3) & (e4 = op(e0, op(op(e0, op(e0, e0)), e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax125)).
fof(f2036, plain, (~ spl24_266 | ~ spl24_64 | ~ spl24_234 | ~ spl24_253), inference(avatar_split_clause, [], [f684, f1893, f1788, f957, f2000])).
fof(f684, plain, (~ (e0 = op(e2, op(e2, e2))) | ~ (e1 = op(op(e2, op(e2, e2)), e2)) | ~ (e3 = op(e2, e2)) | ~ (e4 = op(e2, op(op(e2, op(e2, e2)), e2)))), inference(cnf_transformation, [], [f244])).
fof(f244, plain, (~ (e0 = op(e2, op(e2, e2))) | ~ (e1 = op(op(e2, op(e2, e2)), e2)) | ~ (e3 = op(e2, e2)) | ~ (e4 = op(e2, op(op(e2, op(e2, e2)), e2)))), inference(ennf_transformation, [], [f124])).
fof(f124, plain, ~ ((e0 = op(e2, op(e2, e2))) & (e1 = op(op(e2, op(e2, e2)), e2)) & (e3 = op(e2, e2)) & (e4 = op(e2, op(op(e2, op(e2, e2)), e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax124)).
fof(f2035, plain, (~ spl24_269 | ~ spl24_124 | ~ spl24_259 | ~ spl24_252), inference(avatar_split_clause, [], [f683, f1888, f1938, f1209, f2020])).
fof(f683, plain, (~ (e2 = op(e0, op(e0, e0))) | ~ (e1 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e3) | ~ (e4 = op(e0, op(op(e0, op(e0, e0)), e0)))), inference(cnf_transformation, [], [f243])).
fof(f243, plain, (~ (e2 = op(e0, op(e0, e0))) | ~ (e1 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e3) | ~ (e4 = op(e0, op(op(e0, op(e0, e0)), e0)))), inference(ennf_transformation, [], [f123])).
fof(f123, plain, ~ ((e2 = op(e0, op(e0, e0))) & (e1 = op(op(e0, op(e0, e0)), e0)) & (op(e0, e0) = e3) & (e4 = op(e0, op(op(e0, op(e0, e0)), e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax123)).
fof(f2031, plain, (~ spl24_269 | ~ spl24_123 | ~ spl24_244 | ~ spl24_260), inference(avatar_split_clause, [], [f679, f1944, f1848, f1205, f2020])).
fof(f679, plain, (~ (e1 = op(e0, op(e0, e0))) | ~ (e3 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e2) | ~ (e4 = op(e0, op(op(e0, op(e0, e0)), e0)))), inference(cnf_transformation, [], [f239])).
fof(f239, plain, (~ (e1 = op(e0, op(e0, e0))) | ~ (e3 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e2) | ~ (e4 = op(e0, op(op(e0, op(e0, e0)), e0)))), inference(ennf_transformation, [], [f119])).
fof(f119, plain, ~ ((e1 = op(e0, op(e0, e0))) & (e3 = op(op(e0, op(e0, e0)), e0)) & (op(e0, e0) = e2) & (e4 = op(e0, op(op(e0, op(e0, e0)), e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax119)).
fof(f2030, plain, (~ spl24_267 | ~ spl24_33 | ~ spl24_225 | ~ spl24_249), inference(avatar_split_clause, [], [f678, f1871, f1746, f827, f2005])).
fof(f678, plain, (~ (e0 = op(e3, op(e3, e3))) | ~ (e1 = op(op(e3, op(e3, e3)), e3)) | ~ (e2 = op(e3, e3)) | ~ (e4 = op(e3, op(op(e3, op(e3, e3)), e3)))), inference(cnf_transformation, [], [f238])).
fof(f238, plain, (~ (e0 = op(e3, op(e3, e3))) | ~ (e1 = op(op(e3, op(e3, e3)), e3)) | ~ (e2 = op(e3, e3)) | ~ (e4 = op(e3, op(op(e3, op(e3, e3)), e3)))), inference(ennf_transformation, [], [f118])).
fof(f118, plain, ~ ((e0 = op(e3, op(e3, e3))) & (e1 = op(op(e3, op(e3, e3)), e3)) & (e2 = op(e3, e3)) & (e4 = op(e3, op(op(e3, op(e3, e3)), e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax118)).
fof(f2028, plain, (~ spl24_267 | ~ spl24_33 | ~ spl24_241 | ~ spl24_233), inference(avatar_split_clause, [], [f676, f1783, f1834, f827, f2005])).
fof(f676, plain, (~ (e1 = op(e3, op(e3, e3))) | ~ (e0 = op(op(e3, op(e3, e3)), e3)) | ~ (e2 = op(e3, e3)) | ~ (e4 = op(e3, op(op(e3, op(e3, e3)), e3)))), inference(cnf_transformation, [], [f236])).
fof(f236, plain, (~ (e1 = op(e3, op(e3, e3))) | ~ (e0 = op(op(e3, op(e3, e3)), e3)) | ~ (e2 = op(e3, e3)) | ~ (e4 = op(e3, op(op(e3, op(e3, e3)), e3)))), inference(ennf_transformation, [], [f116])).
fof(f116, plain, ~ ((e1 = op(e3, op(e3, e3))) & (e0 = op(op(e3, op(e3, e3)), e3)) & (e2 = op(e3, e3)) & (e4 = op(e3, op(op(e3, op(e3, e3)), e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax116)).
fof(f2026, plain, (~ spl24_266 | ~ spl24_62 | ~ spl24_217 | ~ spl24_253), inference(avatar_split_clause, [], [f674, f1893, f1710, f949, f2000])).
fof(f674, plain, (~ (e0 = op(e2, op(e2, e2))) | ~ (e3 = op(op(e2, op(e2, e2)), e2)) | ~ (e1 = op(e2, e2)) | ~ (e4 = op(e2, op(op(e2, op(e2, e2)), e2)))), inference(cnf_transformation, [], [f234])).
fof(f234, plain, (~ (e0 = op(e2, op(e2, e2))) | ~ (e3 = op(op(e2, op(e2, e2)), e2)) | ~ (e1 = op(e2, e2)) | ~ (e4 = op(e2, op(op(e2, op(e2, e2)), e2)))), inference(ennf_transformation, [], [f114])).
fof(f114, plain, ~ ((e0 = op(e2, op(e2, e2))) & (e3 = op(op(e2, op(e2, e2)), e2)) & (e1 = op(e2, e2)) & (e4 = op(e2, op(op(e2, op(e2, e2)), e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax114)).
fof(f2023, plain, (~ spl24_269 | ~ spl24_122 | ~ spl24_251 | ~ spl24_248), inference(avatar_split_clause, [], [f671, f1866, f1882, f1201, f2020])).
fof(f671, plain, (~ (e3 = op(e0, op(e0, e0))) | ~ (e2 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e1) | ~ (e4 = op(e0, op(op(e0, op(e0, e0)), e0)))), inference(cnf_transformation, [], [f231])).
fof(f231, plain, (~ (e3 = op(e0, op(e0, e0))) | ~ (e2 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e1) | ~ (e4 = op(e0, op(op(e0, op(e0, e0)), e0)))), inference(ennf_transformation, [], [f111])).
fof(f111, plain, ~ ((e3 = op(e0, op(e0, e0))) & (e2 = op(op(e0, op(e0, e0)), e0)) & (op(e0, e0) = e1) & (e4 = op(e0, op(op(e0, op(e0, e0)), e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax111)).
fof(f2015, plain, (~ spl24_268 | ~ spl24_91 | ~ spl24_228 | ~ spl24_236), inference(avatar_split_clause, [], [f667, f1800, f1760, f1071, f2010])).
fof(f667, plain, (~ (e2 = op(e1, op(e1, e1))) | ~ (e3 = op(op(e1, op(e1, e1)), e1)) | ~ (e0 = op(e1, e1)) | ~ (e4 = op(e1, op(op(e1, op(e1, e1)), e1)))), inference(cnf_transformation, [], [f227])).
fof(f227, plain, (~ (e2 = op(e1, op(e1, e1))) | ~ (e3 = op(op(e1, op(e1, e1)), e1)) | ~ (e0 = op(e1, e1)) | ~ (e4 = op(e1, op(op(e1, op(e1, e1)), e1)))), inference(ennf_transformation, [], [f107])).
fof(f107, plain, ~ ((e2 = op(e1, op(e1, e1))) & (e3 = op(op(e1, op(e1, e1)), e1)) & (e0 = op(e1, e1)) & (e4 = op(e1, op(op(e1, op(e1, e1)), e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax107)).
fof(f2013, plain, (~ spl24_268 | ~ spl24_91 | ~ spl24_235 | ~ spl24_232), inference(avatar_split_clause, [], [f665, f1778, f1794, f1071, f2010])).
fof(f665, plain, (~ (e3 = op(e1, op(e1, e1))) | ~ (e2 = op(op(e1, op(e1, e1)), e1)) | ~ (e0 = op(e1, e1)) | ~ (e4 = op(e1, op(op(e1, op(e1, e1)), e1)))), inference(cnf_transformation, [], [f225])).
fof(f225, plain, (~ (e3 = op(e1, op(e1, e1))) | ~ (e2 = op(op(e1, op(e1, e1)), e1)) | ~ (e0 = op(e1, e1)) | ~ (e4 = op(e1, op(op(e1, op(e1, e1)), e1)))), inference(ennf_transformation, [], [f105])).
fof(f105, plain, ~ ((e3 = op(e1, op(e1, e1))) & (e2 = op(op(e1, op(e1, e1)), e1)) & (e0 = op(e1, e1)) & (e4 = op(e1, op(op(e1, op(e1, e1)), e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax105)).
fof(f2008, plain, (~ spl24_267 | ~ spl24_31 | ~ spl24_225 | ~ spl24_224), inference(avatar_split_clause, [], [f664, f1741, f1746, f819, f2005])).
fof(f664, plain, (~ (e2 = op(e3, op(e3, e3))) | ~ (e1 = op(op(e3, op(e3, e3)), e3)) | ~ (e0 = op(e3, e3)) | ~ (e4 = op(e3, op(op(e3, op(e3, e3)), e3)))), inference(cnf_transformation, [], [f224])).
fof(f224, plain, (~ (e2 = op(e3, op(e3, e3))) | ~ (e1 = op(op(e3, op(e3, e3)), e3)) | ~ (e0 = op(e3, e3)) | ~ (e4 = op(e3, op(op(e3, op(e3, e3)), e3)))), inference(ennf_transformation, [], [f104])).
fof(f104, plain, ~ ((e2 = op(e3, op(e3, e3))) & (e1 = op(op(e3, op(e3, e3)), e3)) & (e0 = op(e3, e3)) & (e4 = op(e3, op(op(e3, op(e3, e3)), e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax104)).
fof(f1997, plain, (~ spl24_265 | ~ spl24_125 | ~ spl24_251 | ~ spl24_260), inference(avatar_split_clause, [], [f661, f1944, f1882, f1213, f1980])).
fof(f661, plain, (~ (e1 = op(e0, op(e0, e0))) | ~ (e2 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e4) | ~ (e3 = op(e0, op(op(e0, op(e0, e0)), e0)))), inference(cnf_transformation, [], [f221])).
fof(f221, plain, (~ (e1 = op(e0, op(e0, e0))) | ~ (e2 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e4) | ~ (e3 = op(e0, op(op(e0, op(e0, e0)), e0)))), inference(ennf_transformation, [], [f101])).
fof(f101, plain, ~ ((e1 = op(e0, op(e0, e0))) & (e2 = op(op(e0, op(e0, e0)), e0)) & (op(e0, e0) = e4) & (e3 = op(e0, op(op(e0, op(e0, e0)), e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax101)).
fof(f1996, plain, (~ spl24_262 | ~ spl24_65 | ~ spl24_234 | ~ spl24_253), inference(avatar_split_clause, [], [f660, f1893, f1788, f961, f1960])).
fof(f660, plain, (~ (e0 = op(e2, op(e2, e2))) | ~ (e1 = op(op(e2, op(e2, e2)), e2)) | ~ (e4 = op(e2, e2)) | ~ (e3 = op(e2, op(op(e2, op(e2, e2)), e2)))), inference(cnf_transformation, [], [f220])).
fof(f220, plain, (~ (e0 = op(e2, op(e2, e2))) | ~ (e1 = op(op(e2, op(e2, e2)), e2)) | ~ (e4 = op(e2, e2)) | ~ (e3 = op(e2, op(op(e2, op(e2, e2)), e2)))), inference(ennf_transformation, [], [f100])).
fof(f100, plain, ~ ((e0 = op(e2, op(e2, e2))) & (e1 = op(op(e2, op(e2, e2)), e2)) & (e4 = op(e2, e2)) & (e3 = op(e2, op(op(e2, op(e2, e2)), e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax100)).
fof(f1995, plain, (~ spl24_265 | ~ spl24_125 | ~ spl24_259 | ~ spl24_252), inference(avatar_split_clause, [], [f659, f1888, f1938, f1213, f1980])).
fof(f659, plain, (~ (e2 = op(e0, op(e0, e0))) | ~ (e1 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e4) | ~ (e3 = op(e0, op(op(e0, op(e0, e0)), e0)))), inference(cnf_transformation, [], [f219])).
fof(f219, plain, (~ (e2 = op(e0, op(e0, e0))) | ~ (e1 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e4) | ~ (e3 = op(e0, op(op(e0, op(e0, e0)), e0)))), inference(ennf_transformation, [], [f99])).
fof(f99, plain, ~ ((e2 = op(e0, op(e0, e0))) & (e1 = op(op(e0, op(e0, e0)), e0)) & (op(e0, e0) = e4) & (e3 = op(e0, op(op(e0, op(e0, e0)), e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax99)).
fof(f1994, plain, (~ spl24_262 | ~ spl24_65 | ~ spl24_250 | ~ spl24_237), inference(avatar_split_clause, [], [f658, f1805, f1876, f961, f1960])).
fof(f658, plain, (~ (e1 = op(e2, op(e2, e2))) | ~ (e0 = op(op(e2, op(e2, e2)), e2)) | ~ (e4 = op(e2, e2)) | ~ (e3 = op(e2, op(op(e2, op(e2, e2)), e2)))), inference(cnf_transformation, [], [f218])).
fof(f218, plain, (~ (e1 = op(e2, op(e2, e2))) | ~ (e0 = op(op(e2, op(e2, e2)), e2)) | ~ (e4 = op(e2, e2)) | ~ (e3 = op(e2, op(op(e2, op(e2, e2)), e2)))), inference(ennf_transformation, [], [f98])).
fof(f98, plain, ~ ((e1 = op(e2, op(e2, e2))) & (e0 = op(op(e2, op(e2, e2)), e2)) & (e4 = op(e2, e2)) & (e3 = op(e2, op(op(e2, op(e2, e2)), e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax98)).
fof(f1993, plain, (~ spl24_264 | ~ spl24_95 | ~ spl24_258 | ~ spl24_236), inference(avatar_split_clause, [], [f657, f1800, f1932, f1087, f1970])).
fof(f657, plain, (~ (e2 = op(e1, op(e1, e1))) | ~ (e0 = op(op(e1, op(e1, e1)), e1)) | ~ (e4 = op(e1, e1)) | ~ (e3 = op(e1, op(op(e1, op(e1, e1)), e1)))), inference(cnf_transformation, [], [f217])).
fof(f217, plain, (~ (e2 = op(e1, op(e1, e1))) | ~ (e0 = op(op(e1, op(e1, e1)), e1)) | ~ (e4 = op(e1, e1)) | ~ (e3 = op(e1, op(op(e1, op(e1, e1)), e1)))), inference(ennf_transformation, [], [f97])).
fof(f97, plain, ~ ((e2 = op(e1, op(e1, e1))) & (e0 = op(op(e1, op(e1, e1)), e1)) & (e4 = op(e1, e1)) & (e3 = op(e1, op(op(e1, op(e1, e1)), e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax97)).
fof(f1992, plain, (~ spl24_264 | ~ spl24_93 | ~ spl24_231 | ~ spl24_261), inference(avatar_split_clause, [], [f656, f1949, f1774, f1079, f1970])).
fof(f656, plain, (~ (e0 = op(e1, op(e1, e1))) | ~ (e4 = op(op(e1, op(e1, e1)), e1)) | ~ (e2 = op(e1, e1)) | ~ (e3 = op(e1, op(op(e1, op(e1, e1)), e1)))), inference(cnf_transformation, [], [f216])).
fof(f216, plain, (~ (e0 = op(e1, op(e1, e1))) | ~ (e4 = op(op(e1, op(e1, e1)), e1)) | ~ (e2 = op(e1, e1)) | ~ (e3 = op(e1, op(op(e1, op(e1, e1)), e1)))), inference(ennf_transformation, [], [f96])).
fof(f96, plain, ~ ((e0 = op(e1, op(e1, e1))) & (e4 = op(op(e1, op(e1, e1)), e1)) & (e2 = op(e1, e1)) & (e3 = op(e1, op(op(e1, op(e1, e1)), e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax96)).
fof(f1990, plain, (~ spl24_263 | ~ spl24_3 | ~ spl24_226 | ~ spl24_246), inference(avatar_split_clause, [], [f654, f1857, f1751, f701, f1965])).
fof(f654, plain, (~ (e0 = op(e4, op(e4, e4))) | ~ (e1 = op(op(e4, op(e4, e4)), e4)) | ~ (e2 = op(e4, e4)) | ~ (e3 = op(e4, op(op(e4, op(e4, e4)), e4)))), inference(cnf_transformation, [], [f214])).
fof(f214, plain, (~ (e0 = op(e4, op(e4, e4))) | ~ (e1 = op(op(e4, op(e4, e4)), e4)) | ~ (e2 = op(e4, e4)) | ~ (e3 = op(e4, op(op(e4, op(e4, e4)), e4)))), inference(ennf_transformation, [], [f94])).
fof(f94, plain, ~ ((e0 = op(e4, op(e4, e4))) & (e1 = op(op(e4, op(e4, e4)), e4)) & (e2 = op(e4, e4)) & (e3 = op(e4, op(op(e4, op(e4, e4)), e4)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax94)).
fof(f1986, plain, (~ spl24_262 | ~ spl24_62 | ~ spl24_221 | ~ spl24_253), inference(avatar_split_clause, [], [f650, f1893, f1728, f949, f1960])).
fof(f650, plain, (~ (e0 = op(e2, op(e2, e2))) | ~ (e4 = op(op(e2, op(e2, e2)), e2)) | ~ (e1 = op(e2, e2)) | ~ (e3 = op(e2, op(op(e2, op(e2, e2)), e2)))), inference(cnf_transformation, [], [f210])).
fof(f210, plain, (~ (e0 = op(e2, op(e2, e2))) | ~ (e4 = op(op(e2, op(e2, e2)), e2)) | ~ (e1 = op(e2, e2)) | ~ (e3 = op(e2, op(op(e2, op(e2, e2)), e2)))), inference(ennf_transformation, [], [f90])).
fof(f90, plain, ~ ((e0 = op(e2, op(e2, e2))) & (e4 = op(op(e2, op(e2, e2)), e2)) & (e1 = op(e2, e2)) & (e3 = op(e2, op(op(e2, op(e2, e2)), e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax90)).
fof(f1985, plain, (~ spl24_265 | ~ spl24_122 | ~ spl24_247 | ~ spl24_252), inference(avatar_split_clause, [], [f649, f1888, f1862, f1201, f1980])).
fof(f649, plain, (~ (e2 = op(e0, op(e0, e0))) | ~ (e4 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e1) | ~ (e3 = op(e0, op(op(e0, op(e0, e0)), e0)))), inference(cnf_transformation, [], [f209])).
fof(f209, plain, (~ (e2 = op(e0, op(e0, e0))) | ~ (e4 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e1) | ~ (e3 = op(e0, op(op(e0, op(e0, e0)), e0)))), inference(ennf_transformation, [], [f89])).
fof(f89, plain, ~ ((e2 = op(e0, op(e0, e0))) & (e4 = op(op(e0, op(e0, e0)), e0)) & (op(e0, e0) = e1) & (e3 = op(e0, op(op(e0, op(e0, e0)), e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax89)).
fof(f1977, plain, (~ spl24_262 | ~ spl24_62 | ~ spl24_250 | ~ spl24_218), inference(avatar_split_clause, [], [f645, f1714, f1876, f949, f1960])).
fof(f645, plain, (~ (e4 = op(e2, op(e2, e2))) | ~ (e0 = op(op(e2, op(e2, e2)), e2)) | ~ (e1 = op(e2, e2)) | ~ (e3 = op(e2, op(op(e2, op(e2, e2)), e2)))), inference(cnf_transformation, [], [f205])).
fof(f205, plain, (~ (e4 = op(e2, op(e2, e2))) | ~ (e0 = op(op(e2, op(e2, e2)), e2)) | ~ (e1 = op(e2, e2)) | ~ (e3 = op(e2, op(op(e2, op(e2, e2)), e2)))), inference(ennf_transformation, [], [f85])).
fof(f85, plain, ~ ((e4 = op(e2, op(e2, e2))) & (e0 = op(op(e2, op(e2, e2)), e2)) & (e1 = op(e2, e2)) & (e3 = op(e2, op(op(e2, op(e2, e2)), e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax85)).
fof(f1975, plain, (~ spl24_264 | ~ spl24_91 | ~ spl24_231 | ~ spl24_236), inference(avatar_split_clause, [], [f643, f1800, f1774, f1071, f1970])).
fof(f643, plain, (~ (e2 = op(e1, op(e1, e1))) | ~ (e4 = op(op(e1, op(e1, e1)), e1)) | ~ (e0 = op(e1, e1)) | ~ (e3 = op(e1, op(op(e1, op(e1, e1)), e1)))), inference(cnf_transformation, [], [f203])).
fof(f203, plain, (~ (e2 = op(e1, op(e1, e1))) | ~ (e4 = op(op(e1, op(e1, e1)), e1)) | ~ (e0 = op(e1, e1)) | ~ (e3 = op(e1, op(op(e1, op(e1, e1)), e1)))), inference(ennf_transformation, [], [f83])).
fof(f83, plain, ~ ((e2 = op(e1, op(e1, e1))) & (e4 = op(op(e1, op(e1, e1)), e1)) & (e0 = op(e1, e1)) & (e3 = op(e1, op(op(e1, op(e1, e1)), e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax83)).
fof(f1963, plain, (~ spl24_262 | ~ spl24_61 | ~ spl24_234 | ~ spl24_218), inference(avatar_split_clause, [], [f639, f1714, f1788, f945, f1960])).
fof(f639, plain, (~ (e4 = op(e2, op(e2, e2))) | ~ (e1 = op(op(e2, op(e2, e2)), e2)) | ~ (e0 = op(e2, e2)) | ~ (e3 = op(e2, op(op(e2, op(e2, e2)), e2)))), inference(cnf_transformation, [], [f199])).
fof(f199, plain, (~ (e4 = op(e2, op(e2, e2))) | ~ (e1 = op(op(e2, op(e2, e2)), e2)) | ~ (e0 = op(e2, e2)) | ~ (e3 = op(e2, op(op(e2, op(e2, e2)), e2)))), inference(ennf_transformation, [], [f79])).
fof(f79, plain, ~ ((e4 = op(e2, op(e2, e2))) & (e1 = op(op(e2, op(e2, e2)), e2)) & (e0 = op(e2, e2)) & (e3 = op(e2, op(op(e2, op(e2, e2)), e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax79)).
fof(f1957, plain, (~ spl24_257 | ~ spl24_125 | ~ spl24_244 | ~ spl24_260), inference(avatar_split_clause, [], [f637, f1944, f1848, f1213, f1924])).
fof(f637, plain, (~ (e1 = op(e0, op(e0, e0))) | ~ (e3 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e4) | ~ (e2 = op(e0, op(op(e0, op(e0, e0)), e0)))), inference(cnf_transformation, [], [f197])).
fof(f197, plain, (~ (e1 = op(e0, op(e0, e0))) | ~ (e3 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e4) | ~ (e2 = op(e0, op(op(e0, op(e0, e0)), e0)))), inference(ennf_transformation, [], [f77])).
fof(f77, plain, ~ ((e1 = op(e0, op(e0, e0))) & (e3 = op(op(e0, op(e0, e0)), e0)) & (op(e0, e0) = e4) & (e2 = op(e0, op(op(e0, op(e0, e0)), e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax77)).
fof(f1956, plain, (~ spl24_254 | ~ spl24_35 | ~ spl24_225 | ~ spl24_249), inference(avatar_split_clause, [], [f636, f1871, f1746, f835, f1904])).
fof(f636, plain, (~ (e0 = op(e3, op(e3, e3))) | ~ (e1 = op(op(e3, op(e3, e3)), e3)) | ~ (e4 = op(e3, e3)) | ~ (e2 = op(e3, op(op(e3, op(e3, e3)), e3)))), inference(cnf_transformation, [], [f196])).
fof(f196, plain, (~ (e0 = op(e3, op(e3, e3))) | ~ (e1 = op(op(e3, op(e3, e3)), e3)) | ~ (e4 = op(e3, e3)) | ~ (e2 = op(e3, op(op(e3, op(e3, e3)), e3)))), inference(ennf_transformation, [], [f76])).
fof(f76, plain, ~ ((e0 = op(e3, op(e3, e3))) & (e1 = op(op(e3, op(e3, e3)), e3)) & (e4 = op(e3, e3)) & (e2 = op(e3, op(op(e3, op(e3, e3)), e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax76)).
fof(f1955, plain, (~ spl24_257 | ~ spl24_125 | ~ spl24_259 | ~ spl24_248), inference(avatar_split_clause, [], [f635, f1866, f1938, f1213, f1924])).
fof(f635, plain, (~ (e3 = op(e0, op(e0, e0))) | ~ (e1 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e4) | ~ (e2 = op(e0, op(op(e0, op(e0, e0)), e0)))), inference(cnf_transformation, [], [f195])).
fof(f195, plain, (~ (e3 = op(e0, op(e0, e0))) | ~ (e1 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e4) | ~ (e2 = op(e0, op(op(e0, op(e0, e0)), e0)))), inference(ennf_transformation, [], [f75])).
fof(f75, plain, ~ ((e3 = op(e0, op(e0, e0))) & (e1 = op(op(e0, op(e0, e0)), e0)) & (op(e0, e0) = e4) & (e2 = op(e0, op(op(e0, op(e0, e0)), e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax75)).
fof(f1952, plain, (~ spl24_256 | ~ spl24_94 | ~ spl24_231 | ~ spl24_261), inference(avatar_split_clause, [], [f632, f1949, f1774, f1083, f1914])).
fof(f632, plain, (~ (e0 = op(e1, op(e1, e1))) | ~ (e4 = op(op(e1, op(e1, e1)), e1)) | ~ (e3 = op(e1, e1)) | ~ (e2 = op(e1, op(op(e1, op(e1, e1)), e1)))), inference(cnf_transformation, [], [f192])).
fof(f192, plain, (~ (e0 = op(e1, op(e1, e1))) | ~ (e4 = op(op(e1, op(e1, e1)), e1)) | ~ (e3 = op(e1, e1)) | ~ (e2 = op(e1, op(op(e1, op(e1, e1)), e1)))), inference(ennf_transformation, [], [f72])).
fof(f72, plain, ~ ((e0 = op(e1, op(e1, e1))) & (e4 = op(op(e1, op(e1, e1)), e1)) & (e3 = op(e1, e1)) & (e2 = op(e1, op(op(e1, op(e1, e1)), e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax72)).
fof(f1947, plain, (~ spl24_257 | ~ spl24_124 | ~ spl24_247 | ~ spl24_260), inference(avatar_split_clause, [], [f631, f1944, f1862, f1209, f1924])).
fof(f631, plain, (~ (e1 = op(e0, op(e0, e0))) | ~ (e4 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e3) | ~ (e2 = op(e0, op(op(e0, op(e0, e0)), e0)))), inference(cnf_transformation, [], [f191])).
fof(f191, plain, (~ (e1 = op(e0, op(e0, e0))) | ~ (e4 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e3) | ~ (e2 = op(e0, op(op(e0, op(e0, e0)), e0)))), inference(ennf_transformation, [], [f71])).
fof(f71, plain, ~ ((e1 = op(e0, op(e0, e0))) & (e4 = op(op(e0, op(e0, e0)), e0)) & (op(e0, e0) = e3) & (e2 = op(e0, op(op(e0, op(e0, e0)), e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax71)).
fof(f1941, plain, (~ spl24_257 | ~ spl24_124 | ~ spl24_259 | ~ spl24_245), inference(avatar_split_clause, [], [f629, f1852, f1938, f1209, f1924])).
fof(f629, plain, (~ (e4 = op(e0, op(e0, e0))) | ~ (e1 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e3) | ~ (e2 = op(e0, op(op(e0, op(e0, e0)), e0)))), inference(cnf_transformation, [], [f189])).
fof(f189, plain, (~ (e4 = op(e0, op(e0, e0))) | ~ (e1 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e3) | ~ (e2 = op(e0, op(op(e0, op(e0, e0)), e0)))), inference(ennf_transformation, [], [f69])).
fof(f69, plain, ~ ((e4 = op(e0, op(e0, e0))) & (e1 = op(op(e0, op(e0, e0)), e0)) & (op(e0, e0) = e3) & (e2 = op(e0, op(op(e0, op(e0, e0)), e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax69)).
fof(f1930, plain, (~ spl24_254 | ~ spl24_32 | ~ spl24_223 | ~ spl24_249), inference(avatar_split_clause, [], [f626, f1871, f1737, f823, f1904])).
fof(f626, plain, (~ (e0 = op(e3, op(e3, e3))) | ~ (e4 = op(op(e3, op(e3, e3)), e3)) | ~ (e1 = op(e3, e3)) | ~ (e2 = op(e3, op(op(e3, op(e3, e3)), e3)))), inference(cnf_transformation, [], [f186])).
fof(f186, plain, (~ (e0 = op(e3, op(e3, e3))) | ~ (e4 = op(op(e3, op(e3, e3)), e3)) | ~ (e1 = op(e3, e3)) | ~ (e2 = op(e3, op(op(e3, op(e3, e3)), e3)))), inference(ennf_transformation, [], [f66])).
fof(f66, plain, ~ ((e0 = op(e3, op(e3, e3))) & (e4 = op(op(e3, op(e3, e3)), e3)) & (e1 = op(e3, e3)) & (e2 = op(e3, op(op(e3, op(e3, e3)), e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax66)).
fof(f1922, plain, (~ spl24_255 | ~ spl24_2 | ~ spl24_242 | ~ spl24_215), inference(avatar_split_clause, [], [f622, f1701, f1839, f697, f1909])).
fof(f622, plain, (~ (e3 = op(e4, op(e4, e4))) | ~ (e0 = op(op(e4, op(e4, e4)), e4)) | ~ (e1 = op(e4, e4)) | ~ (e2 = op(e4, op(op(e4, op(e4, e4)), e4)))), inference(cnf_transformation, [], [f182])).
fof(f182, plain, (~ (e3 = op(e4, op(e4, e4))) | ~ (e0 = op(op(e4, op(e4, e4)), e4)) | ~ (e1 = op(e4, e4)) | ~ (e2 = op(e4, op(op(e4, op(e4, e4)), e4)))), inference(ennf_transformation, [], [f62])).
fof(f62, plain, ~ ((e3 = op(e4, op(e4, e4))) & (e0 = op(op(e4, op(e4, e4)), e4)) & (e1 = op(e4, e4)) & (e2 = op(e4, op(op(e4, op(e4, e4)), e4)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax62)).
fof(f1920, plain, (~ spl24_254 | ~ spl24_31 | ~ spl24_223 | ~ spl24_233), inference(avatar_split_clause, [], [f620, f1783, f1737, f819, f1904])).
fof(f620, plain, (~ (e1 = op(e3, op(e3, e3))) | ~ (e4 = op(op(e3, op(e3, e3)), e3)) | ~ (e0 = op(e3, e3)) | ~ (e2 = op(e3, op(op(e3, op(e3, e3)), e3)))), inference(cnf_transformation, [], [f180])).
fof(f180, plain, (~ (e1 = op(e3, op(e3, e3))) | ~ (e4 = op(op(e3, op(e3, e3)), e3)) | ~ (e0 = op(e3, e3)) | ~ (e2 = op(e3, op(op(e3, op(e3, e3)), e3)))), inference(ennf_transformation, [], [f60])).
fof(f60, plain, ~ ((e1 = op(e3, op(e3, e3))) & (e4 = op(op(e3, op(e3, e3)), e3)) & (e0 = op(e3, e3)) & (e2 = op(e3, op(op(e3, op(e3, e3)), e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax60)).
fof(f1917, plain, (~ spl24_256 | ~ spl24_91 | ~ spl24_228 | ~ spl24_229), inference(avatar_split_clause, [], [f617, f1764, f1760, f1071, f1914])).
fof(f617, plain, (~ (e4 = op(e1, op(e1, e1))) | ~ (e3 = op(op(e1, op(e1, e1)), e1)) | ~ (e0 = op(e1, e1)) | ~ (e2 = op(e1, op(op(e1, op(e1, e1)), e1)))), inference(cnf_transformation, [], [f177])).
fof(f177, plain, (~ (e4 = op(e1, op(e1, e1))) | ~ (e3 = op(op(e1, op(e1, e1)), e1)) | ~ (e0 = op(e1, e1)) | ~ (e2 = op(e1, op(op(e1, op(e1, e1)), e1)))), inference(ennf_transformation, [], [f57])).
fof(f57, plain, ~ ((e4 = op(e1, op(e1, e1))) & (e3 = op(op(e1, op(e1, e1)), e1)) & (e0 = op(e1, e1)) & (e2 = op(e1, op(op(e1, op(e1, e1)), e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax57)).
fof(f1912, plain, (~ spl24_255 | ~ spl24_1 | ~ spl24_226 | ~ spl24_215), inference(avatar_split_clause, [], [f616, f1701, f1751, f693, f1909])).
fof(f616, plain, (~ (e3 = op(e4, op(e4, e4))) | ~ (e1 = op(op(e4, op(e4, e4)), e4)) | ~ (e0 = op(e4, e4)) | ~ (e2 = op(e4, op(op(e4, op(e4, e4)), e4)))), inference(cnf_transformation, [], [f176])).
fof(f176, plain, (~ (e3 = op(e4, op(e4, e4))) | ~ (e1 = op(op(e4, op(e4, e4)), e4)) | ~ (e0 = op(e4, e4)) | ~ (e2 = op(e4, op(op(e4, op(e4, e4)), e4)))), inference(ennf_transformation, [], [f56])).
fof(f56, plain, ~ ((e3 = op(e4, op(e4, e4))) & (e1 = op(op(e4, op(e4, e4)), e4)) & (e0 = op(e4, e4)) & (e2 = op(e4, op(op(e4, op(e4, e4)), e4)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax56)).
fof(f1907, plain, (~ spl24_254 | ~ spl24_31 | ~ spl24_225 | ~ spl24_212), inference(avatar_split_clause, [], [f615, f1688, f1746, f819, f1904])).
fof(f615, plain, (~ (e4 = op(e3, op(e3, e3))) | ~ (e1 = op(op(e3, op(e3, e3)), e3)) | ~ (e0 = op(e3, e3)) | ~ (e2 = op(e3, op(op(e3, op(e3, e3)), e3)))), inference(cnf_transformation, [], [f175])).
fof(f175, plain, (~ (e4 = op(e3, op(e3, e3))) | ~ (e1 = op(op(e3, op(e3, e3)), e3)) | ~ (e0 = op(e3, e3)) | ~ (e2 = op(e3, op(op(e3, op(e3, e3)), e3)))), inference(ennf_transformation, [], [f55])).
fof(f55, plain, ~ ((e4 = op(e3, op(e3, e3))) & (e1 = op(op(e3, op(e3, e3)), e3)) & (e0 = op(e3, e3)) & (e2 = op(e3, op(op(e3, op(e3, e3)), e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax55)).
fof(f1901, plain, (~ spl24_243 | ~ spl24_125 | ~ spl24_244 | ~ spl24_252), inference(avatar_split_clause, [], [f613, f1888, f1848, f1213, f1844])).
fof(f613, plain, (~ (e2 = op(e0, op(e0, e0))) | ~ (e3 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e4) | ~ (e1 = op(e0, op(op(e0, op(e0, e0)), e0)))), inference(cnf_transformation, [], [f173])).
fof(f173, plain, (~ (e2 = op(e0, op(e0, e0))) | ~ (e3 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e4) | ~ (e1 = op(e0, op(op(e0, op(e0, e0)), e0)))), inference(ennf_transformation, [], [f53])).
fof(f53, plain, ~ ((e2 = op(e0, op(e0, e0))) & (e3 = op(op(e0, op(e0, e0)), e0)) & (op(e0, e0) = e4) & (e1 = op(e0, op(op(e0, op(e0, e0)), e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax53)).
fof(f1899, plain, (~ spl24_243 | ~ spl24_125 | ~ spl24_251 | ~ spl24_248), inference(avatar_split_clause, [], [f611, f1866, f1882, f1213, f1844])).
fof(f611, plain, (~ (e3 = op(e0, op(e0, e0))) | ~ (e2 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e4) | ~ (e1 = op(e0, op(op(e0, op(e0, e0)), e0)))), inference(cnf_transformation, [], [f171])).
fof(f171, plain, (~ (e3 = op(e0, op(e0, e0))) | ~ (e2 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e4) | ~ (e1 = op(e0, op(op(e0, op(e0, e0)), e0)))), inference(ennf_transformation, [], [f51])).
fof(f51, plain, ~ ((e3 = op(e0, op(e0, e0))) & (e2 = op(op(e0, op(e0, e0)), e0)) & (op(e0, e0) = e4) & (e1 = op(e0, op(op(e0, op(e0, e0)), e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax51)).
fof(f1898, plain, (~ spl24_238 | ~ spl24_35 | ~ spl24_241 | ~ spl24_224), inference(avatar_split_clause, [], [f610, f1741, f1834, f835, f1816])).
fof(f610, plain, (~ (e2 = op(e3, op(e3, e3))) | ~ (e0 = op(op(e3, op(e3, e3)), e3)) | ~ (e4 = op(e3, e3)) | ~ (e1 = op(e3, op(op(e3, op(e3, e3)), e3)))), inference(cnf_transformation, [], [f170])).
fof(f170, plain, (~ (e2 = op(e3, op(e3, e3))) | ~ (e0 = op(op(e3, op(e3, e3)), e3)) | ~ (e4 = op(e3, e3)) | ~ (e1 = op(e3, op(op(e3, op(e3, e3)), e3)))), inference(ennf_transformation, [], [f50])).
fof(f50, plain, ~ ((e2 = op(e3, op(e3, e3))) & (e0 = op(op(e3, op(e3, e3)), e3)) & (e4 = op(e3, e3)) & (e1 = op(e3, op(op(e3, op(e3, e3)), e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax50)).
fof(f1897, plain, (~ spl24_240 | ~ spl24_65 | ~ spl24_250 | ~ spl24_222), inference(avatar_split_clause, [], [f609, f1732, f1876, f961, f1826])).
fof(f609, plain, (~ (e3 = op(e2, op(e2, e2))) | ~ (e0 = op(op(e2, op(e2, e2)), e2)) | ~ (e4 = op(e2, e2)) | ~ (e1 = op(e2, op(op(e2, op(e2, e2)), e2)))), inference(cnf_transformation, [], [f169])).
fof(f169, plain, (~ (e3 = op(e2, op(e2, e2))) | ~ (e0 = op(op(e2, op(e2, e2)), e2)) | ~ (e4 = op(e2, e2)) | ~ (e1 = op(e2, op(op(e2, op(e2, e2)), e2)))), inference(ennf_transformation, [], [f49])).
fof(f49, plain, ~ ((e3 = op(e2, op(e2, e2))) & (e0 = op(op(e2, op(e2, e2)), e2)) & (e4 = op(e2, e2)) & (e1 = op(e2, op(op(e2, op(e2, e2)), e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax49)).
fof(f1896, plain, (~ spl24_240 | ~ spl24_64 | ~ spl24_221 | ~ spl24_253), inference(avatar_split_clause, [], [f608, f1893, f1728, f957, f1826])).
fof(f608, plain, (~ (e0 = op(e2, op(e2, e2))) | ~ (e4 = op(op(e2, op(e2, e2)), e2)) | ~ (e3 = op(e2, e2)) | ~ (e1 = op(e2, op(op(e2, op(e2, e2)), e2)))), inference(cnf_transformation, [], [f168])).
fof(f168, plain, (~ (e0 = op(e2, op(e2, e2))) | ~ (e4 = op(op(e2, op(e2, e2)), e2)) | ~ (e3 = op(e2, e2)) | ~ (e1 = op(e2, op(op(e2, op(e2, e2)), e2)))), inference(ennf_transformation, [], [f48])).
fof(f48, plain, ~ ((e0 = op(e2, op(e2, e2))) & (e4 = op(op(e2, op(e2, e2)), e2)) & (e3 = op(e2, e2)) & (e1 = op(e2, op(op(e2, op(e2, e2)), e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax48)).
fof(f1891, plain, (~ spl24_243 | ~ spl24_124 | ~ spl24_247 | ~ spl24_252), inference(avatar_split_clause, [], [f607, f1888, f1862, f1209, f1844])).
fof(f607, plain, (~ (e2 = op(e0, op(e0, e0))) | ~ (e4 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e3) | ~ (e1 = op(e0, op(op(e0, op(e0, e0)), e0)))), inference(cnf_transformation, [], [f167])).
fof(f167, plain, (~ (e2 = op(e0, op(e0, e0))) | ~ (e4 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e3) | ~ (e1 = op(e0, op(op(e0, op(e0, e0)), e0)))), inference(ennf_transformation, [], [f47])).
fof(f47, plain, ~ ((e2 = op(e0, op(e0, e0))) & (e4 = op(op(e0, op(e0, e0)), e0)) & (op(e0, e0) = e3) & (e1 = op(e0, op(op(e0, op(e0, e0)), e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax47)).
fof(f1885, plain, (~ spl24_243 | ~ spl24_124 | ~ spl24_251 | ~ spl24_245), inference(avatar_split_clause, [], [f605, f1852, f1882, f1209, f1844])).
fof(f605, plain, (~ (e4 = op(e0, op(e0, e0))) | ~ (e2 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e3) | ~ (e1 = op(e0, op(op(e0, op(e0, e0)), e0)))), inference(cnf_transformation, [], [f165])).
fof(f165, plain, (~ (e4 = op(e0, op(e0, e0))) | ~ (e2 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e3) | ~ (e1 = op(e0, op(op(e0, op(e0, e0)), e0)))), inference(ennf_transformation, [], [f45])).
fof(f45, plain, ~ ((e4 = op(e0, op(e0, e0))) & (e2 = op(op(e0, op(e0, e0)), e0)) & (op(e0, e0) = e3) & (e1 = op(e0, op(op(e0, op(e0, e0)), e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax45)).
fof(f1874, plain, (~ spl24_238 | ~ spl24_33 | ~ spl24_223 | ~ spl24_249), inference(avatar_split_clause, [], [f602, f1871, f1737, f827, f1816])).
fof(f602, plain, (~ (e0 = op(e3, op(e3, e3))) | ~ (e4 = op(op(e3, op(e3, e3)), e3)) | ~ (e2 = op(e3, e3)) | ~ (e1 = op(e3, op(op(e3, op(e3, e3)), e3)))), inference(cnf_transformation, [], [f162])).
fof(f162, plain, (~ (e0 = op(e3, op(e3, e3))) | ~ (e4 = op(op(e3, op(e3, e3)), e3)) | ~ (e2 = op(e3, e3)) | ~ (e1 = op(e3, op(op(e3, op(e3, e3)), e3)))), inference(ennf_transformation, [], [f42])).
fof(f42, plain, ~ ((e0 = op(e3, op(e3, e3))) & (e4 = op(op(e3, op(e3, e3)), e3)) & (e2 = op(e3, e3)) & (e1 = op(e3, op(op(e3, op(e3, e3)), e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax42)).
fof(f1837, plain, (~ spl24_238 | ~ spl24_33 | ~ spl24_241 | ~ spl24_212), inference(avatar_split_clause, [], [f597, f1688, f1834, f827, f1816])).
fof(f597, plain, (~ (e4 = op(e3, op(e3, e3))) | ~ (e0 = op(op(e3, op(e3, e3)), e3)) | ~ (e2 = op(e3, e3)) | ~ (e1 = op(e3, op(op(e3, op(e3, e3)), e3)))), inference(cnf_transformation, [], [f157])).
fof(f157, plain, (~ (e4 = op(e3, op(e3, e3))) | ~ (e0 = op(op(e3, op(e3, e3)), e3)) | ~ (e2 = op(e3, e3)) | ~ (e1 = op(e3, op(op(e3, op(e3, e3)), e3)))), inference(ennf_transformation, [], [f37])).
fof(f37, plain, ~ ((e4 = op(e3, op(e3, e3))) & (e0 = op(op(e3, op(e3, e3)), e3)) & (e2 = op(e3, e3)) & (e1 = op(e3, op(op(e3, op(e3, e3)), e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax37)).
fof(f1819, plain, (~ spl24_238 | ~ spl24_31 | ~ spl24_211 | ~ spl24_212), inference(avatar_split_clause, [], [f591, f1688, f1684, f819, f1816])).
fof(f591, plain, (~ (e4 = op(e3, op(e3, e3))) | ~ (e2 = op(op(e3, op(e3, e3)), e3)) | ~ (e0 = op(e3, e3)) | ~ (e1 = op(e3, op(op(e3, op(e3, e3)), e3)))), inference(cnf_transformation, [], [f151])).
fof(f151, plain, (~ (e4 = op(e3, op(e3, e3))) | ~ (e2 = op(op(e3, op(e3, e3)), e3)) | ~ (e0 = op(e3, e3)) | ~ (e1 = op(e3, op(op(e3, op(e3, e3)), e3)))), inference(ennf_transformation, [], [f31])).
fof(f31, plain, ~ ((e4 = op(e3, op(e3, e3))) & (e2 = op(op(e3, op(e3, e3)), e3)) & (e0 = op(e3, e3)) & (e1 = op(e3, op(op(e3, op(e3, e3)), e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax31)).
fof(f1813, plain, (~ spl24_227 | ~ spl24_95 | ~ spl24_228 | ~ spl24_236), inference(avatar_split_clause, [], [f589, f1800, f1760, f1087, f1756])).
fof(f589, plain, (~ (e2 = op(e1, op(e1, e1))) | ~ (e3 = op(op(e1, op(e1, e1)), e1)) | ~ (e4 = op(e1, e1)) | ~ (e0 = op(e1, op(op(e1, op(e1, e1)), e1)))), inference(cnf_transformation, [], [f149])).
fof(f149, plain, (~ (e2 = op(e1, op(e1, e1))) | ~ (e3 = op(op(e1, op(e1, e1)), e1)) | ~ (e4 = op(e1, e1)) | ~ (e0 = op(e1, op(op(e1, op(e1, e1)), e1)))), inference(ennf_transformation, [], [f29])).
fof(f29, plain, ~ ((e2 = op(e1, op(e1, e1))) & (e3 = op(op(e1, op(e1, e1)), e1)) & (e4 = op(e1, e1)) & (e0 = op(e1, op(op(e1, op(e1, e1)), e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax29)).
fof(f1811, plain, (~ spl24_227 | ~ spl24_95 | ~ spl24_235 | ~ spl24_232), inference(avatar_split_clause, [], [f587, f1778, f1794, f1087, f1756])).
fof(f587, plain, (~ (e3 = op(e1, op(e1, e1))) | ~ (e2 = op(op(e1, op(e1, e1)), e1)) | ~ (e4 = op(e1, e1)) | ~ (e0 = op(e1, op(op(e1, op(e1, e1)), e1)))), inference(cnf_transformation, [], [f147])).
fof(f147, plain, (~ (e3 = op(e1, op(e1, e1))) | ~ (e2 = op(op(e1, op(e1, e1)), e1)) | ~ (e4 = op(e1, e1)) | ~ (e0 = op(e1, op(op(e1, op(e1, e1)), e1)))), inference(ennf_transformation, [], [f27])).
fof(f27, plain, ~ ((e3 = op(e1, op(e1, e1))) & (e2 = op(op(e1, op(e1, e1)), e1)) & (e4 = op(e1, e1)) & (e0 = op(e1, op(op(e1, op(e1, e1)), e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax27)).
fof(f1810, plain, (~ spl24_210 | ~ spl24_35 | ~ spl24_225 | ~ spl24_224), inference(avatar_split_clause, [], [f586, f1741, f1746, f835, f1680])).
fof(f586, plain, (~ (e2 = op(e3, op(e3, e3))) | ~ (e1 = op(op(e3, op(e3, e3)), e3)) | ~ (e4 = op(e3, e3)) | ~ (e0 = op(e3, op(op(e3, op(e3, e3)), e3)))), inference(cnf_transformation, [], [f146])).
fof(f146, plain, (~ (e2 = op(e3, op(e3, e3))) | ~ (e1 = op(op(e3, op(e3, e3)), e3)) | ~ (e4 = op(e3, e3)) | ~ (e0 = op(e3, op(op(e3, op(e3, e3)), e3)))), inference(ennf_transformation, [], [f26])).
fof(f26, plain, ~ ((e2 = op(e3, op(e3, e3))) & (e1 = op(op(e3, op(e3, e3)), e3)) & (e4 = op(e3, e3)) & (e0 = op(e3, op(op(e3, op(e3, e3)), e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax26)).
fof(f1798, plain, (~ spl24_213 | ~ spl24_4 | ~ spl24_214 | ~ spl24_230), inference(avatar_split_clause, [], [f582, f1769, f1697, f705, f1693])).
fof(f582, plain, (~ (e1 = op(e4, op(e4, e4))) | ~ (e2 = op(op(e4, op(e4, e4)), e4)) | ~ (e3 = op(e4, e4)) | ~ (e0 = op(e4, op(op(e4, op(e4, e4)), e4)))), inference(cnf_transformation, [], [f142])).
fof(f142, plain, (~ (e1 = op(e4, op(e4, e4))) | ~ (e2 = op(op(e4, op(e4, e4)), e4)) | ~ (e3 = op(e4, e4)) | ~ (e0 = op(e4, op(op(e4, op(e4, e4)), e4)))), inference(ennf_transformation, [], [f22])).
fof(f22, plain, ~ ((e1 = op(e4, op(e4, e4))) & (e2 = op(op(e4, op(e4, e4)), e4)) & (e3 = op(e4, e4)) & (e0 = op(e4, op(op(e4, op(e4, e4)), e4)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax22)).
fof(f1797, plain, (~ spl24_227 | ~ spl24_94 | ~ spl24_235 | ~ spl24_229), inference(avatar_split_clause, [], [f581, f1764, f1794, f1083, f1756])).
fof(f581, plain, (~ (e4 = op(e1, op(e1, e1))) | ~ (e2 = op(op(e1, op(e1, e1)), e1)) | ~ (e3 = op(e1, e1)) | ~ (e0 = op(e1, op(op(e1, op(e1, e1)), e1)))), inference(cnf_transformation, [], [f141])).
fof(f141, plain, (~ (e4 = op(e1, op(e1, e1))) | ~ (e2 = op(op(e1, op(e1, e1)), e1)) | ~ (e3 = op(e1, e1)) | ~ (e0 = op(e1, op(op(e1, op(e1, e1)), e1)))), inference(ennf_transformation, [], [f21])).
fof(f21, plain, ~ ((e4 = op(e1, op(e1, e1))) & (e2 = op(op(e1, op(e1, e1)), e1)) & (e3 = op(e1, e1)) & (e0 = op(e1, op(op(e1, op(e1, e1)), e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax21)).
fof(f1786, plain, (~ spl24_210 | ~ spl24_33 | ~ spl24_223 | ~ spl24_233), inference(avatar_split_clause, [], [f578, f1783, f1737, f827, f1680])).
fof(f578, plain, (~ (e1 = op(e3, op(e3, e3))) | ~ (e4 = op(op(e3, op(e3, e3)), e3)) | ~ (e2 = op(e3, e3)) | ~ (e0 = op(e3, op(op(e3, op(e3, e3)), e3)))), inference(cnf_transformation, [], [f138])).
fof(f138, plain, (~ (e1 = op(e3, op(e3, e3))) | ~ (e4 = op(op(e3, op(e3, e3)), e3)) | ~ (e2 = op(e3, e3)) | ~ (e0 = op(e3, op(op(e3, op(e3, e3)), e3)))), inference(ennf_transformation, [], [f18])).
fof(f18, plain, ~ ((e1 = op(e3, op(e3, e3))) & (e4 = op(op(e3, op(e3, e3)), e3)) & (e2 = op(e3, e3)) & (e0 = op(e3, op(op(e3, op(e3, e3)), e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax18)).
fof(f1781, plain, (~ spl24_227 | ~ spl24_93 | ~ spl24_231 | ~ spl24_232), inference(avatar_split_clause, [], [f577, f1778, f1774, f1079, f1756])).
fof(f577, plain, (~ (e3 = op(e1, op(e1, e1))) | ~ (e4 = op(op(e1, op(e1, e1)), e1)) | ~ (e2 = op(e1, e1)) | ~ (e0 = op(e1, op(op(e1, op(e1, e1)), e1)))), inference(cnf_transformation, [], [f137])).
fof(f137, plain, (~ (e3 = op(e1, op(e1, e1))) | ~ (e4 = op(op(e1, op(e1, e1)), e1)) | ~ (e2 = op(e1, e1)) | ~ (e0 = op(e1, op(op(e1, op(e1, e1)), e1)))), inference(ennf_transformation, [], [f17])).
fof(f17, plain, ~ ((e3 = op(e1, op(e1, e1))) & (e4 = op(op(e1, op(e1, e1)), e1)) & (e2 = op(e1, e1)) & (e0 = op(e1, op(op(e1, op(e1, e1)), e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax17)).
fof(f1772, plain, (~ spl24_213 | ~ spl24_3 | ~ spl24_219 | ~ spl24_230), inference(avatar_split_clause, [], [f576, f1769, f1719, f701, f1693])).
fof(f576, plain, (~ (e1 = op(e4, op(e4, e4))) | ~ (e3 = op(op(e4, op(e4, e4)), e4)) | ~ (e2 = op(e4, e4)) | ~ (e0 = op(e4, op(op(e4, op(e4, e4)), e4)))), inference(cnf_transformation, [], [f136])).
fof(f136, plain, (~ (e1 = op(e4, op(e4, e4))) | ~ (e3 = op(op(e4, op(e4, e4)), e4)) | ~ (e2 = op(e4, e4)) | ~ (e0 = op(e4, op(op(e4, op(e4, e4)), e4)))), inference(ennf_transformation, [], [f16])).
fof(f16, plain, ~ ((e1 = op(e4, op(e4, e4))) & (e3 = op(op(e4, op(e4, e4)), e4)) & (e2 = op(e4, e4)) & (e0 = op(e4, op(op(e4, op(e4, e4)), e4)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax16)).
fof(f1767, plain, (~ spl24_227 | ~ spl24_93 | ~ spl24_228 | ~ spl24_229), inference(avatar_split_clause, [], [f575, f1764, f1760, f1079, f1756])).
fof(f575, plain, (~ (e4 = op(e1, op(e1, e1))) | ~ (e3 = op(op(e1, op(e1, e1)), e1)) | ~ (e2 = op(e1, e1)) | ~ (e0 = op(e1, op(op(e1, op(e1, e1)), e1)))), inference(cnf_transformation, [], [f135])).
fof(f135, plain, (~ (e4 = op(e1, op(e1, e1))) | ~ (e3 = op(op(e1, op(e1, e1)), e1)) | ~ (e2 = op(e1, e1)) | ~ (e0 = op(e1, op(op(e1, op(e1, e1)), e1)))), inference(ennf_transformation, [], [f15])).
fof(f15, plain, ~ ((e4 = op(e1, op(e1, e1))) & (e3 = op(op(e1, op(e1, e1)), e1)) & (e2 = op(e1, e1)) & (e0 = op(e1, op(op(e1, op(e1, e1)), e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax15)).
fof(f1749, plain, (~ spl24_210 | ~ spl24_33 | ~ spl24_225 | ~ spl24_212), inference(avatar_split_clause, [], [f573, f1688, f1746, f827, f1680])).
fof(f573, plain, (~ (e4 = op(e3, op(e3, e3))) | ~ (e1 = op(op(e3, op(e3, e3)), e3)) | ~ (e2 = op(e3, e3)) | ~ (e0 = op(e3, op(op(e3, op(e3, e3)), e3)))), inference(cnf_transformation, [], [f133])).
fof(f133, plain, (~ (e4 = op(e3, op(e3, e3))) | ~ (e1 = op(op(e3, op(e3, e3)), e3)) | ~ (e2 = op(e3, e3)) | ~ (e0 = op(e3, op(op(e3, op(e3, e3)), e3)))), inference(ennf_transformation, [], [f13])).
fof(f13, plain, ~ ((e4 = op(e3, op(e3, e3))) & (e1 = op(op(e3, op(e3, e3)), e3)) & (e2 = op(e3, e3)) & (e0 = op(e3, op(op(e3, op(e3, e3)), e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax13)).
fof(f1744, plain, (~ spl24_210 | ~ spl24_32 | ~ spl24_223 | ~ spl24_224), inference(avatar_split_clause, [], [f572, f1741, f1737, f823, f1680])).
fof(f572, plain, (~ (e2 = op(e3, op(e3, e3))) | ~ (e4 = op(op(e3, op(e3, e3)), e3)) | ~ (e1 = op(e3, e3)) | ~ (e0 = op(e3, op(op(e3, op(e3, e3)), e3)))), inference(cnf_transformation, [], [f132])).
fof(f132, plain, (~ (e2 = op(e3, op(e3, e3))) | ~ (e4 = op(op(e3, op(e3, e3)), e3)) | ~ (e1 = op(e3, e3)) | ~ (e0 = op(e3, op(op(e3, op(e3, e3)), e3)))), inference(ennf_transformation, [], [f12])).
fof(f12, plain, ~ ((e2 = op(e3, op(e3, e3))) & (e4 = op(op(e3, op(e3, e3)), e3)) & (e1 = op(e3, e3)) & (e0 = op(e3, op(op(e3, op(e3, e3)), e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax12)).
fof(f1735, plain, (~ spl24_216 | ~ spl24_62 | ~ spl24_221 | ~ spl24_222), inference(avatar_split_clause, [], [f571, f1732, f1728, f949, f1706])).
fof(f571, plain, (~ (e3 = op(e2, op(e2, e2))) | ~ (e4 = op(op(e2, op(e2, e2)), e2)) | ~ (e1 = op(e2, e2)) | ~ (e0 = op(e2, op(op(e2, op(e2, e2)), e2)))), inference(cnf_transformation, [], [f131])).
fof(f131, plain, (~ (e3 = op(e2, op(e2, e2))) | ~ (e4 = op(op(e2, op(e2, e2)), e2)) | ~ (e1 = op(e2, e2)) | ~ (e0 = op(e2, op(op(e2, op(e2, e2)), e2)))), inference(ennf_transformation, [], [f11])).
fof(f11, plain, ~ ((e3 = op(e2, op(e2, e2))) & (e4 = op(op(e2, op(e2, e2)), e2)) & (e1 = op(e2, e2)) & (e0 = op(e2, op(op(e2, op(e2, e2)), e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax11)).
fof(f1726, plain, (~ spl24_213 | ~ spl24_2 | ~ spl24_219 | ~ spl24_220), inference(avatar_split_clause, [], [f570, f1723, f1719, f697, f1693])).
fof(f570, plain, (~ (e2 = op(e4, op(e4, e4))) | ~ (e3 = op(op(e4, op(e4, e4)), e4)) | ~ (e1 = op(e4, e4)) | ~ (e0 = op(e4, op(op(e4, op(e4, e4)), e4)))), inference(cnf_transformation, [], [f130])).
fof(f130, plain, (~ (e2 = op(e4, op(e4, e4))) | ~ (e3 = op(op(e4, op(e4, e4)), e4)) | ~ (e1 = op(e4, e4)) | ~ (e0 = op(e4, op(op(e4, op(e4, e4)), e4)))), inference(ennf_transformation, [], [f10])).
fof(f10, plain, ~ ((e2 = op(e4, op(e4, e4))) & (e3 = op(op(e4, op(e4, e4)), e4)) & (e1 = op(e4, e4)) & (e0 = op(e4, op(op(e4, op(e4, e4)), e4)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax10)).
fof(f1717, plain, (~ spl24_216 | ~ spl24_62 | ~ spl24_217 | ~ spl24_218), inference(avatar_split_clause, [], [f569, f1714, f1710, f949, f1706])).
fof(f569, plain, (~ (e4 = op(e2, op(e2, e2))) | ~ (e3 = op(op(e2, op(e2, e2)), e2)) | ~ (e1 = op(e2, e2)) | ~ (e0 = op(e2, op(op(e2, op(e2, e2)), e2)))), inference(cnf_transformation, [], [f129])).
fof(f129, plain, (~ (e4 = op(e2, op(e2, e2))) | ~ (e3 = op(op(e2, op(e2, e2)), e2)) | ~ (e1 = op(e2, e2)) | ~ (e0 = op(e2, op(op(e2, op(e2, e2)), e2)))), inference(ennf_transformation, [], [f9])).
fof(f9, plain, ~ ((e4 = op(e2, op(e2, e2))) & (e3 = op(op(e2, op(e2, e2)), e2)) & (e1 = op(e2, e2)) & (e0 = op(e2, op(op(e2, op(e2, e2)), e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax9)).
fof(f1691, plain, (~ spl24_210 | ~ spl24_32 | ~ spl24_211 | ~ spl24_212), inference(avatar_split_clause, [], [f567, f1688, f1684, f823, f1680])).
fof(f567, plain, (~ (e4 = op(e3, op(e3, e3))) | ~ (e2 = op(op(e3, op(e3, e3)), e3)) | ~ (e1 = op(e3, e3)) | ~ (e0 = op(e3, op(op(e3, op(e3, e3)), e3)))), inference(cnf_transformation, [], [f127])).
fof(f127, plain, (~ (e4 = op(e3, op(e3, e3))) | ~ (e2 = op(op(e3, op(e3, e3)), e3)) | ~ (e1 = op(e3, e3)) | ~ (e0 = op(e3, op(op(e3, op(e3, e3)), e3)))), inference(ennf_transformation, [], [f7])).
fof(f7, plain, ~ ((e4 = op(e3, op(e3, e3))) & (e2 = op(op(e3, op(e3, e3)), e3)) & (e1 = op(e3, e3)) & (e0 = op(e3, op(op(e3, op(e3, e3)), e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax7)).
fof(f1678, plain, (spl24_207 | spl24_204 | spl24_201 | spl24_198 | spl24_195 | spl24_191 | spl24_189 | spl24_186 | spl24_183 | spl24_180 | spl24_176 | spl24_172 | spl24_170 | spl24_167 | spl24_164 | spl24_160 | spl24_156 | spl24_152 | spl24_150 | spl24_147 | spl24_143 | spl24_139 | spl24_135 | spl24_131), inference(avatar_split_clause, [], [f687, f1289, f1308, f1327, f1346, f1365, f1380, f1391, f1410, f1429, f1448, f1463, f1478, f1489, f1508, f1527, f1542, f1557, f1572, f1583, f1602, f1617, f1632, f1647, f1662])).
fof(f1662, plain, (spl24_207 <=> sP0), introduced(avatar_definition, [new_symbols(naming, [spl24_207])])).
fof(f1647, plain, (spl24_204 <=> sP1), introduced(avatar_definition, [new_symbols(naming, [spl24_204])])).
fof(f1632, plain, (spl24_201 <=> sP2), introduced(avatar_definition, [new_symbols(naming, [spl24_201])])).
fof(f1617, plain, (spl24_198 <=> sP3), introduced(avatar_definition, [new_symbols(naming, [spl24_198])])).
fof(f1602, plain, (spl24_195 <=> sP4), introduced(avatar_definition, [new_symbols(naming, [spl24_195])])).
fof(f1583, plain, (spl24_191 <=> sP5), introduced(avatar_definition, [new_symbols(naming, [spl24_191])])).
fof(f1572, plain, (spl24_189 <=> sP6), introduced(avatar_definition, [new_symbols(naming, [spl24_189])])).
fof(f1557, plain, (spl24_186 <=> sP7), introduced(avatar_definition, [new_symbols(naming, [spl24_186])])).
fof(f1542, plain, (spl24_183 <=> sP8), introduced(avatar_definition, [new_symbols(naming, [spl24_183])])).
fof(f1527, plain, (spl24_180 <=> sP9), introduced(avatar_definition, [new_symbols(naming, [spl24_180])])).
fof(f1508, plain, (spl24_176 <=> sP10), introduced(avatar_definition, [new_symbols(naming, [spl24_176])])).
fof(f1489, plain, (spl24_172 <=> sP11), introduced(avatar_definition, [new_symbols(naming, [spl24_172])])).
fof(f1478, plain, (spl24_170 <=> sP12), introduced(avatar_definition, [new_symbols(naming, [spl24_170])])).
fof(f1463, plain, (spl24_167 <=> sP13), introduced(avatar_definition, [new_symbols(naming, [spl24_167])])).
fof(f1448, plain, (spl24_164 <=> sP14), introduced(avatar_definition, [new_symbols(naming, [spl24_164])])).
fof(f1429, plain, (spl24_160 <=> sP15), introduced(avatar_definition, [new_symbols(naming, [spl24_160])])).
fof(f1410, plain, (spl24_156 <=> sP16), introduced(avatar_definition, [new_symbols(naming, [spl24_156])])).
fof(f1391, plain, (spl24_152 <=> sP17), introduced(avatar_definition, [new_symbols(naming, [spl24_152])])).
fof(f1380, plain, (spl24_150 <=> sP18), introduced(avatar_definition, [new_symbols(naming, [spl24_150])])).
fof(f1365, plain, (spl24_147 <=> sP19), introduced(avatar_definition, [new_symbols(naming, [spl24_147])])).
fof(f1346, plain, (spl24_143 <=> sP20), introduced(avatar_definition, [new_symbols(naming, [spl24_143])])).
fof(f1327, plain, (spl24_139 <=> sP21), introduced(avatar_definition, [new_symbols(naming, [spl24_139])])).
fof(f1308, plain, (spl24_135 <=> sP22), introduced(avatar_definition, [new_symbols(naming, [spl24_135])])).
fof(f1289, plain, (spl24_131 <=> sP23), introduced(avatar_definition, [new_symbols(naming, [spl24_131])])).
fof(f687, plain, (sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(trivial_inequality_removal, [], [f564])).
fof(f564, plain, (~ (op(e4, e4) = op(e4, e4)) | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f271])).
fof(f271, plain, ((~ (e4 = op(op(e4, e4), e4)) & (e4 = op(op(e4, e4), e4)) & ~ (op(e4, e4) = op(e4, e4))) | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(definition_folding, [], [f6, e270, e269, e268, e267, e266, e265, e264, e263, e262, e261, e260, e259, e258, e257, e256, e255, e254, e253, e252, e251, e250, e249, e248, e247])).
fof(f247, plain, ((~ (e0 = op(op(e0, e0), e0)) & (e0 = op(op(e0, e0), e0)) & ~ (op(e0, e0) = op(e0, e0))) | ~ sP0), inference(usedef, [], [e247])).
fof(e247, plain, (sP0 <=> (~ (e0 = op(op(e0, e0), e0)) & (e0 = op(op(e0, e0), e0)) & ~ (op(e0, e0) = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f248, plain, ((~ (e1 = op(op(e0, e1), e0)) & (e0 = op(op(e0, e1), e1)) & ~ (op(e0, e1) = op(e1, e0))) | ~ sP1), inference(usedef, [], [e248])).
fof(e248, plain, (sP1 <=> (~ (e1 = op(op(e0, e1), e0)) & (e0 = op(op(e0, e1), e1)) & ~ (op(e0, e1) = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f249, plain, ((~ (e2 = op(op(e0, e2), e0)) & (e0 = op(op(e0, e2), e2)) & ~ (op(e0, e2) = op(e2, e0))) | ~ sP2), inference(usedef, [], [e249])).
fof(e249, plain, (sP2 <=> (~ (e2 = op(op(e0, e2), e0)) & (e0 = op(op(e0, e2), e2)) & ~ (op(e0, e2) = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f250, plain, ((~ (e3 = op(op(e0, e3), e0)) & (e0 = op(op(e0, e3), e3)) & ~ (op(e0, e3) = op(e3, e0))) | ~ sP3), inference(usedef, [], [e250])).
fof(e250, plain, (sP3 <=> (~ (e3 = op(op(e0, e3), e0)) & (e0 = op(op(e0, e3), e3)) & ~ (op(e0, e3) = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f251, plain, ((~ (e4 = op(op(e0, e4), e0)) & (e0 = op(op(e0, e4), e4)) & ~ (op(e0, e4) = op(e4, e0))) | ~ sP4), inference(usedef, [], [e251])).
fof(e251, plain, (sP4 <=> (~ (e4 = op(op(e0, e4), e0)) & (e0 = op(op(e0, e4), e4)) & ~ (op(e0, e4) = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f252, plain, ((~ (e0 = op(op(e1, e0), e1)) & (e1 = op(op(e1, e0), e0)) & ~ (op(e0, e1) = op(e1, e0))) | ~ sP5), inference(usedef, [], [e252])).
fof(e252, plain, (sP5 <=> (~ (e0 = op(op(e1, e0), e1)) & (e1 = op(op(e1, e0), e0)) & ~ (op(e0, e1) = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f253, plain, ((~ (e1 = op(op(e1, e1), e1)) & (e1 = op(op(e1, e1), e1)) & ~ (op(e1, e1) = op(e1, e1))) | ~ sP6), inference(usedef, [], [e253])).
fof(e253, plain, (sP6 <=> (~ (e1 = op(op(e1, e1), e1)) & (e1 = op(op(e1, e1), e1)) & ~ (op(e1, e1) = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f254, plain, ((~ (e2 = op(op(e1, e2), e1)) & (e1 = op(op(e1, e2), e2)) & ~ (op(e1, e2) = op(e2, e1))) | ~ sP7), inference(usedef, [], [e254])).
fof(e254, plain, (sP7 <=> (~ (e2 = op(op(e1, e2), e1)) & (e1 = op(op(e1, e2), e2)) & ~ (op(e1, e2) = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f255, plain, ((~ (e3 = op(op(e1, e3), e1)) & (e1 = op(op(e1, e3), e3)) & ~ (op(e1, e3) = op(e3, e1))) | ~ sP8), inference(usedef, [], [e255])).
fof(e255, plain, (sP8 <=> (~ (e3 = op(op(e1, e3), e1)) & (e1 = op(op(e1, e3), e3)) & ~ (op(e1, e3) = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f256, plain, ((~ (e4 = op(op(e1, e4), e1)) & (e1 = op(op(e1, e4), e4)) & ~ (op(e1, e4) = op(e4, e1))) | ~ sP9), inference(usedef, [], [e256])).
fof(e256, plain, (sP9 <=> (~ (e4 = op(op(e1, e4), e1)) & (e1 = op(op(e1, e4), e4)) & ~ (op(e1, e4) = op(e4, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f257, plain, ((~ (e0 = op(op(e2, e0), e2)) & (e2 = op(op(e2, e0), e0)) & ~ (op(e0, e2) = op(e2, e0))) | ~ sP10), inference(usedef, [], [e257])).
fof(e257, plain, (sP10 <=> (~ (e0 = op(op(e2, e0), e2)) & (e2 = op(op(e2, e0), e0)) & ~ (op(e0, e2) = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f258, plain, ((~ (e1 = op(op(e2, e1), e2)) & (e2 = op(op(e2, e1), e1)) & ~ (op(e1, e2) = op(e2, e1))) | ~ sP11), inference(usedef, [], [e258])).
fof(e258, plain, (sP11 <=> (~ (e1 = op(op(e2, e1), e2)) & (e2 = op(op(e2, e1), e1)) & ~ (op(e1, e2) = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f259, plain, ((~ (e2 = op(op(e2, e2), e2)) & (e2 = op(op(e2, e2), e2)) & ~ (op(e2, e2) = op(e2, e2))) | ~ sP12), inference(usedef, [], [e259])).
fof(e259, plain, (sP12 <=> (~ (e2 = op(op(e2, e2), e2)) & (e2 = op(op(e2, e2), e2)) & ~ (op(e2, e2) = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f260, plain, ((~ (e3 = op(op(e2, e3), e2)) & (e2 = op(op(e2, e3), e3)) & ~ (op(e2, e3) = op(e3, e2))) | ~ sP13), inference(usedef, [], [e260])).
fof(e260, plain, (sP13 <=> (~ (e3 = op(op(e2, e3), e2)) & (e2 = op(op(e2, e3), e3)) & ~ (op(e2, e3) = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f261, plain, ((~ (e4 = op(op(e2, e4), e2)) & (e2 = op(op(e2, e4), e4)) & ~ (op(e2, e4) = op(e4, e2))) | ~ sP14), inference(usedef, [], [e261])).
fof(e261, plain, (sP14 <=> (~ (e4 = op(op(e2, e4), e2)) & (e2 = op(op(e2, e4), e4)) & ~ (op(e2, e4) = op(e4, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f262, plain, ((~ (e0 = op(op(e3, e0), e3)) & (e3 = op(op(e3, e0), e0)) & ~ (op(e0, e3) = op(e3, e0))) | ~ sP15), inference(usedef, [], [e262])).
fof(e262, plain, (sP15 <=> (~ (e0 = op(op(e3, e0), e3)) & (e3 = op(op(e3, e0), e0)) & ~ (op(e0, e3) = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f263, plain, ((~ (e1 = op(op(e3, e1), e3)) & (e3 = op(op(e3, e1), e1)) & ~ (op(e1, e3) = op(e3, e1))) | ~ sP16), inference(usedef, [], [e263])).
fof(e263, plain, (sP16 <=> (~ (e1 = op(op(e3, e1), e3)) & (e3 = op(op(e3, e1), e1)) & ~ (op(e1, e3) = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f264, plain, ((~ (e2 = op(op(e3, e2), e3)) & (e3 = op(op(e3, e2), e2)) & ~ (op(e2, e3) = op(e3, e2))) | ~ sP17), inference(usedef, [], [e264])).
fof(e264, plain, (sP17 <=> (~ (e2 = op(op(e3, e2), e3)) & (e3 = op(op(e3, e2), e2)) & ~ (op(e2, e3) = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f265, plain, ((~ (e3 = op(op(e3, e3), e3)) & (e3 = op(op(e3, e3), e3)) & ~ (op(e3, e3) = op(e3, e3))) | ~ sP18), inference(usedef, [], [e265])).
fof(e265, plain, (sP18 <=> (~ (e3 = op(op(e3, e3), e3)) & (e3 = op(op(e3, e3), e3)) & ~ (op(e3, e3) = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f266, plain, ((~ (e4 = op(op(e3, e4), e3)) & (e3 = op(op(e3, e4), e4)) & ~ (op(e3, e4) = op(e4, e3))) | ~ sP19), inference(usedef, [], [e266])).
fof(e266, plain, (sP19 <=> (~ (e4 = op(op(e3, e4), e3)) & (e3 = op(op(e3, e4), e4)) & ~ (op(e3, e4) = op(e4, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP19])])).
fof(f267, plain, ((~ (e0 = op(op(e4, e0), e4)) & (e4 = op(op(e4, e0), e0)) & ~ (op(e0, e4) = op(e4, e0))) | ~ sP20), inference(usedef, [], [e267])).
fof(e267, plain, (sP20 <=> (~ (e0 = op(op(e4, e0), e4)) & (e4 = op(op(e4, e0), e0)) & ~ (op(e0, e4) = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP20])])).
fof(f268, plain, ((~ (e1 = op(op(e4, e1), e4)) & (e4 = op(op(e4, e1), e1)) & ~ (op(e1, e4) = op(e4, e1))) | ~ sP21), inference(usedef, [], [e268])).
fof(e268, plain, (sP21 <=> (~ (e1 = op(op(e4, e1), e4)) & (e4 = op(op(e4, e1), e1)) & ~ (op(e1, e4) = op(e4, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP21])])).
fof(f269, plain, ((~ (e2 = op(op(e4, e2), e4)) & (e4 = op(op(e4, e2), e2)) & ~ (op(e2, e4) = op(e4, e2))) | ~ sP22), inference(usedef, [], [e269])).
fof(e269, plain, (sP22 <=> (~ (e2 = op(op(e4, e2), e4)) & (e4 = op(op(e4, e2), e2)) & ~ (op(e2, e4) = op(e4, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP22])])).
fof(f270, plain, ((~ (e3 = op(op(e4, e3), e4)) & (e4 = op(op(e4, e3), e3)) & ~ (op(e3, e4) = op(e4, e3))) | ~ sP23), inference(usedef, [], [e270])).
fof(e270, plain, (sP23 <=> (~ (e3 = op(op(e4, e3), e4)) & (e4 = op(op(e4, e3), e3)) & ~ (op(e3, e4) = op(e4, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP23])])).
fof(f6, plain, ((~ (e4 = op(op(e4, e4), e4)) & (e4 = op(op(e4, e4), e4)) & ~ (op(e4, e4) = op(e4, e4))) | (~ (e3 = op(op(e4, e3), e4)) & (e4 = op(op(e4, e3), e3)) & ~ (op(e3, e4) = op(e4, e3))) | (~ (e2 = op(op(e4, e2), e4)) & (e4 = op(op(e4, e2), e2)) & ~ (op(e2, e4) = op(e4, e2))) | (~ (e1 = op(op(e4, e1), e4)) & (e4 = op(op(e4, e1), e1)) & ~ (op(e1, e4) = op(e4, e1))) | (~ (e0 = op(op(e4, e0), e4)) & (e4 = op(op(e4, e0), e0)) & ~ (op(e0, e4) = op(e4, e0))) | (~ (e4 = op(op(e3, e4), e3)) & (e3 = op(op(e3, e4), e4)) & ~ (op(e3, e4) = op(e4, e3))) | (~ (e3 = op(op(e3, e3), e3)) & (e3 = op(op(e3, e3), e3)) & ~ (op(e3, e3) = op(e3, e3))) | (~ (e2 = op(op(e3, e2), e3)) & (e3 = op(op(e3, e2), e2)) & ~ (op(e2, e3) = op(e3, e2))) | (~ (e1 = op(op(e3, e1), e3)) & (e3 = op(op(e3, e1), e1)) & ~ (op(e1, e3) = op(e3, e1))) | (~ (e0 = op(op(e3, e0), e3)) & (e3 = op(op(e3, e0), e0)) & ~ (op(e0, e3) = op(e3, e0))) | (~ (e4 = op(op(e2, e4), e2)) & (e2 = op(op(e2, e4), e4)) & ~ (op(e2, e4) = op(e4, e2))) | (~ (e3 = op(op(e2, e3), e2)) & (e2 = op(op(e2, e3), e3)) & ~ (op(e2, e3) = op(e3, e2))) | (~ (e2 = op(op(e2, e2), e2)) & (e2 = op(op(e2, e2), e2)) & ~ (op(e2, e2) = op(e2, e2))) | (~ (e1 = op(op(e2, e1), e2)) & (e2 = op(op(e2, e1), e1)) & ~ (op(e1, e2) = op(e2, e1))) | (~ (e0 = op(op(e2, e0), e2)) & (e2 = op(op(e2, e0), e0)) & ~ (op(e0, e2) = op(e2, e0))) | (~ (e4 = op(op(e1, e4), e1)) & (e1 = op(op(e1, e4), e4)) & ~ (op(e1, e4) = op(e4, e1))) | (~ (e3 = op(op(e1, e3), e1)) & (e1 = op(op(e1, e3), e3)) & ~ (op(e1, e3) = op(e3, e1))) | (~ (e2 = op(op(e1, e2), e1)) & (e1 = op(op(e1, e2), e2)) & ~ (op(e1, e2) = op(e2, e1))) | (~ (e1 = op(op(e1, e1), e1)) & (e1 = op(op(e1, e1), e1)) & ~ (op(e1, e1) = op(e1, e1))) | (~ (e0 = op(op(e1, e0), e1)) & (e1 = op(op(e1, e0), e0)) & ~ (op(e0, e1) = op(e1, e0))) | (~ (e4 = op(op(e0, e4), e0)) & (e0 = op(op(e0, e4), e4)) & ~ (op(e0, e4) = op(e4, e0))) | (~ (e3 = op(op(e0, e3), e0)) & (e0 = op(op(e0, e3), e3)) & ~ (op(e0, e3) = op(e3, e0))) | (~ (e2 = op(op(e0, e2), e0)) & (e0 = op(op(e0, e2), e2)) & ~ (op(e0, e2) = op(e2, e0))) | (~ (e1 = op(op(e0, e1), e0)) & (e0 = op(op(e0, e1), e1)) & ~ (op(e0, e1) = op(e1, e0))) | (~ (e0 = op(op(e0, e0), e0)) & (e0 = op(op(e0, e0), e0)) & ~ (op(e0, e0) = op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax6)).
fof(f1671, plain, ~ spl24_207, inference(avatar_split_clause, [], [f688, f1662])).
fof(f688, plain, ~ sP0, inference(trivial_inequality_removal, [], [f561])).
fof(f561, plain, (~ (op(e0, e0) = op(e0, e0)) | ~ sP0), inference(cnf_transformation, [], [f295])).
fof(f295, plain, ((~ (e0 = op(op(e0, e0), e0)) & (e0 = op(op(e0, e0), e0)) & ~ (op(e0, e0) = op(e0, e0))) | ~ sP0), inference(nnf_transformation, [], [f247])).
fof(f1660, plain, (~ spl24_204 | ~ spl24_194), inference(avatar_split_clause, [], [f558, f1597, f1647])).
fof(f558, plain, (~ (op(e0, e1) = op(e1, e0)) | ~ sP1), inference(cnf_transformation, [], [f294])).
fof(f294, plain, ((~ (e1 = op(op(e0, e1), e0)) & (e0 = op(op(e0, e1), e1)) & ~ (op(e0, e1) = op(e1, e0))) | ~ sP1), inference(nnf_transformation, [], [f248])).
fof(f1659, plain, (~ spl24_204 | spl24_206), inference(avatar_split_clause, [], [f559, f1656, f1647])).
fof(f559, plain, ((e0 = op(op(e0, e1), e1)) | ~ sP1), inference(cnf_transformation, [], [f294])).
fof(f1654, plain, (~ spl24_204 | ~ spl24_205), inference(avatar_split_clause, [], [f560, f1651, f1647])).
fof(f560, plain, (~ (e1 = op(op(e0, e1), e0)) | ~ sP1), inference(cnf_transformation, [], [f294])).
fof(f1645, plain, (~ spl24_201 | ~ spl24_179), inference(avatar_split_clause, [], [f555, f1522, f1632])).
fof(f555, plain, (~ (op(e0, e2) = op(e2, e0)) | ~ sP2), inference(cnf_transformation, [], [f293])).
fof(f293, plain, ((~ (e2 = op(op(e0, e2), e0)) & (e0 = op(op(e0, e2), e2)) & ~ (op(e0, e2) = op(e2, e0))) | ~ sP2), inference(nnf_transformation, [], [f249])).
fof(f1644, plain, (~ spl24_201 | spl24_203), inference(avatar_split_clause, [], [f556, f1641, f1632])).
fof(f556, plain, ((e0 = op(op(e0, e2), e2)) | ~ sP2), inference(cnf_transformation, [], [f293])).
fof(f1639, plain, (~ spl24_201 | ~ spl24_202), inference(avatar_split_clause, [], [f557, f1636, f1632])).
fof(f557, plain, (~ (e2 = op(op(e0, e2), e0)) | ~ sP2), inference(cnf_transformation, [], [f293])).
fof(f1630, plain, (~ spl24_198 | ~ spl24_163), inference(avatar_split_clause, [], [f552, f1443, f1617])).
fof(f552, plain, (~ (op(e0, e3) = op(e3, e0)) | ~ sP3), inference(cnf_transformation, [], [f292])).
fof(f292, plain, ((~ (e3 = op(op(e0, e3), e0)) & (e0 = op(op(e0, e3), e3)) & ~ (op(e0, e3) = op(e3, e0))) | ~ sP3), inference(nnf_transformation, [], [f250])).
fof(f1629, plain, (~ spl24_198 | spl24_200), inference(avatar_split_clause, [], [f553, f1626, f1617])).
fof(f553, plain, ((e0 = op(op(e0, e3), e3)) | ~ sP3), inference(cnf_transformation, [], [f292])).
fof(f1624, plain, (~ spl24_198 | ~ spl24_199), inference(avatar_split_clause, [], [f554, f1621, f1617])).
fof(f554, plain, (~ (e3 = op(op(e0, e3), e0)) | ~ sP3), inference(cnf_transformation, [], [f292])).
fof(f1615, plain, (~ spl24_195 | ~ spl24_146), inference(avatar_split_clause, [], [f549, f1360, f1602])).
fof(f549, plain, (~ (op(e0, e4) = op(e4, e0)) | ~ sP4), inference(cnf_transformation, [], [f291])).
fof(f291, plain, ((~ (e4 = op(op(e0, e4), e0)) & (e0 = op(op(e0, e4), e4)) & ~ (op(e0, e4) = op(e4, e0))) | ~ sP4), inference(nnf_transformation, [], [f251])).
fof(f1614, plain, (~ spl24_195 | spl24_197), inference(avatar_split_clause, [], [f550, f1611, f1602])).
fof(f550, plain, ((e0 = op(op(e0, e4), e4)) | ~ sP4), inference(cnf_transformation, [], [f291])).
fof(f1609, plain, (~ spl24_195 | ~ spl24_196), inference(avatar_split_clause, [], [f551, f1606, f1602])).
fof(f551, plain, (~ (e4 = op(op(e0, e4), e0)) | ~ sP4), inference(cnf_transformation, [], [f291])).
fof(f1600, plain, (~ spl24_191 | ~ spl24_194), inference(avatar_split_clause, [], [f546, f1597, f1583])).
fof(f546, plain, (~ (op(e0, e1) = op(e1, e0)) | ~ sP5), inference(cnf_transformation, [], [f290])).
fof(f290, plain, ((~ (e0 = op(op(e1, e0), e1)) & (e1 = op(op(e1, e0), e0)) & ~ (op(e0, e1) = op(e1, e0))) | ~ sP5), inference(nnf_transformation, [], [f252])).
fof(f1595, plain, (~ spl24_191 | spl24_193), inference(avatar_split_clause, [], [f547, f1592, f1583])).
fof(f547, plain, ((e1 = op(op(e1, e0), e0)) | ~ sP5), inference(cnf_transformation, [], [f290])).
fof(f1590, plain, (~ spl24_191 | ~ spl24_192), inference(avatar_split_clause, [], [f548, f1587, f1583])).
fof(f548, plain, (~ (e0 = op(op(e1, e0), e1)) | ~ sP5), inference(cnf_transformation, [], [f290])).
fof(f1581, plain, ~ spl24_189, inference(avatar_split_clause, [], [f689, f1572])).
fof(f689, plain, ~ sP6, inference(trivial_inequality_removal, [], [f543])).
fof(f543, plain, (~ (op(e1, e1) = op(e1, e1)) | ~ sP6), inference(cnf_transformation, [], [f289])).
fof(f289, plain, ((~ (e1 = op(op(e1, e1), e1)) & (e1 = op(op(e1, e1), e1)) & ~ (op(e1, e1) = op(e1, e1))) | ~ sP6), inference(nnf_transformation, [], [f253])).
fof(f1570, plain, (~ spl24_186 | ~ spl24_175), inference(avatar_split_clause, [], [f540, f1503, f1557])).
fof(f540, plain, (~ (op(e1, e2) = op(e2, e1)) | ~ sP7), inference(cnf_transformation, [], [f288])).
fof(f288, plain, ((~ (e2 = op(op(e1, e2), e1)) & (e1 = op(op(e1, e2), e2)) & ~ (op(e1, e2) = op(e2, e1))) | ~ sP7), inference(nnf_transformation, [], [f254])).
fof(f1569, plain, (~ spl24_186 | spl24_188), inference(avatar_split_clause, [], [f541, f1566, f1557])).
fof(f541, plain, ((e1 = op(op(e1, e2), e2)) | ~ sP7), inference(cnf_transformation, [], [f288])).
fof(f1564, plain, (~ spl24_186 | ~ spl24_187), inference(avatar_split_clause, [], [f542, f1561, f1557])).
fof(f542, plain, (~ (e2 = op(op(e1, e2), e1)) | ~ sP7), inference(cnf_transformation, [], [f288])).
fof(f1555, plain, (~ spl24_183 | ~ spl24_159), inference(avatar_split_clause, [], [f537, f1424, f1542])).
fof(f537, plain, (~ (op(e1, e3) = op(e3, e1)) | ~ sP8), inference(cnf_transformation, [], [f287])).
fof(f287, plain, ((~ (e3 = op(op(e1, e3), e1)) & (e1 = op(op(e1, e3), e3)) & ~ (op(e1, e3) = op(e3, e1))) | ~ sP8), inference(nnf_transformation, [], [f255])).
fof(f1554, plain, (~ spl24_183 | spl24_185), inference(avatar_split_clause, [], [f538, f1551, f1542])).
fof(f538, plain, ((e1 = op(op(e1, e3), e3)) | ~ sP8), inference(cnf_transformation, [], [f287])).
fof(f1549, plain, (~ spl24_183 | ~ spl24_184), inference(avatar_split_clause, [], [f539, f1546, f1542])).
fof(f539, plain, (~ (e3 = op(op(e1, e3), e1)) | ~ sP8), inference(cnf_transformation, [], [f287])).
fof(f1540, plain, (~ spl24_180 | ~ spl24_142), inference(avatar_split_clause, [], [f534, f1341, f1527])).
fof(f534, plain, (~ (op(e1, e4) = op(e4, e1)) | ~ sP9), inference(cnf_transformation, [], [f286])).
fof(f286, plain, ((~ (e4 = op(op(e1, e4), e1)) & (e1 = op(op(e1, e4), e4)) & ~ (op(e1, e4) = op(e4, e1))) | ~ sP9), inference(nnf_transformation, [], [f256])).
fof(f1539, plain, (~ spl24_180 | spl24_182), inference(avatar_split_clause, [], [f535, f1536, f1527])).
fof(f535, plain, ((e1 = op(op(e1, e4), e4)) | ~ sP9), inference(cnf_transformation, [], [f286])).
fof(f1534, plain, (~ spl24_180 | ~ spl24_181), inference(avatar_split_clause, [], [f536, f1531, f1527])).
fof(f536, plain, (~ (e4 = op(op(e1, e4), e1)) | ~ sP9), inference(cnf_transformation, [], [f286])).
fof(f1525, plain, (~ spl24_176 | ~ spl24_179), inference(avatar_split_clause, [], [f531, f1522, f1508])).
fof(f531, plain, (~ (op(e0, e2) = op(e2, e0)) | ~ sP10), inference(cnf_transformation, [], [f285])).
fof(f285, plain, ((~ (e0 = op(op(e2, e0), e2)) & (e2 = op(op(e2, e0), e0)) & ~ (op(e0, e2) = op(e2, e0))) | ~ sP10), inference(nnf_transformation, [], [f257])).
fof(f1520, plain, (~ spl24_176 | spl24_178), inference(avatar_split_clause, [], [f532, f1517, f1508])).
fof(f532, plain, ((e2 = op(op(e2, e0), e0)) | ~ sP10), inference(cnf_transformation, [], [f285])).
fof(f1515, plain, (~ spl24_176 | ~ spl24_177), inference(avatar_split_clause, [], [f533, f1512, f1508])).
fof(f533, plain, (~ (e0 = op(op(e2, e0), e2)) | ~ sP10), inference(cnf_transformation, [], [f285])).
fof(f1506, plain, (~ spl24_172 | ~ spl24_175), inference(avatar_split_clause, [], [f528, f1503, f1489])).
fof(f528, plain, (~ (op(e1, e2) = op(e2, e1)) | ~ sP11), inference(cnf_transformation, [], [f284])).
fof(f284, plain, ((~ (e1 = op(op(e2, e1), e2)) & (e2 = op(op(e2, e1), e1)) & ~ (op(e1, e2) = op(e2, e1))) | ~ sP11), inference(nnf_transformation, [], [f258])).
fof(f1501, plain, (~ spl24_172 | spl24_174), inference(avatar_split_clause, [], [f529, f1498, f1489])).
fof(f529, plain, ((e2 = op(op(e2, e1), e1)) | ~ sP11), inference(cnf_transformation, [], [f284])).
fof(f1496, plain, (~ spl24_172 | ~ spl24_173), inference(avatar_split_clause, [], [f530, f1493, f1489])).
fof(f530, plain, (~ (e1 = op(op(e2, e1), e2)) | ~ sP11), inference(cnf_transformation, [], [f284])).
fof(f1487, plain, ~ spl24_170, inference(avatar_split_clause, [], [f690, f1478])).
fof(f690, plain, ~ sP12, inference(trivial_inequality_removal, [], [f525])).
fof(f525, plain, (~ (op(e2, e2) = op(e2, e2)) | ~ sP12), inference(cnf_transformation, [], [f283])).
fof(f283, plain, ((~ (e2 = op(op(e2, e2), e2)) & (e2 = op(op(e2, e2), e2)) & ~ (op(e2, e2) = op(e2, e2))) | ~ sP12), inference(nnf_transformation, [], [f259])).
fof(f1476, plain, (~ spl24_167 | ~ spl24_155), inference(avatar_split_clause, [], [f522, f1405, f1463])).
fof(f522, plain, (~ (op(e2, e3) = op(e3, e2)) | ~ sP13), inference(cnf_transformation, [], [f282])).
fof(f282, plain, ((~ (e3 = op(op(e2, e3), e2)) & (e2 = op(op(e2, e3), e3)) & ~ (op(e2, e3) = op(e3, e2))) | ~ sP13), inference(nnf_transformation, [], [f260])).
fof(f1475, plain, (~ spl24_167 | spl24_169), inference(avatar_split_clause, [], [f523, f1472, f1463])).
fof(f523, plain, ((e2 = op(op(e2, e3), e3)) | ~ sP13), inference(cnf_transformation, [], [f282])).
fof(f1470, plain, (~ spl24_167 | ~ spl24_168), inference(avatar_split_clause, [], [f524, f1467, f1463])).
fof(f524, plain, (~ (e3 = op(op(e2, e3), e2)) | ~ sP13), inference(cnf_transformation, [], [f282])).
fof(f1461, plain, (~ spl24_164 | ~ spl24_138), inference(avatar_split_clause, [], [f519, f1322, f1448])).
fof(f519, plain, (~ (op(e2, e4) = op(e4, e2)) | ~ sP14), inference(cnf_transformation, [], [f281])).
fof(f281, plain, ((~ (e4 = op(op(e2, e4), e2)) & (e2 = op(op(e2, e4), e4)) & ~ (op(e2, e4) = op(e4, e2))) | ~ sP14), inference(nnf_transformation, [], [f261])).
fof(f1460, plain, (~ spl24_164 | spl24_166), inference(avatar_split_clause, [], [f520, f1457, f1448])).
fof(f520, plain, ((e2 = op(op(e2, e4), e4)) | ~ sP14), inference(cnf_transformation, [], [f281])).
fof(f1455, plain, (~ spl24_164 | ~ spl24_165), inference(avatar_split_clause, [], [f521, f1452, f1448])).
fof(f521, plain, (~ (e4 = op(op(e2, e4), e2)) | ~ sP14), inference(cnf_transformation, [], [f281])).
fof(f1446, plain, (~ spl24_160 | ~ spl24_163), inference(avatar_split_clause, [], [f516, f1443, f1429])).
fof(f516, plain, (~ (op(e0, e3) = op(e3, e0)) | ~ sP15), inference(cnf_transformation, [], [f280])).
fof(f280, plain, ((~ (e0 = op(op(e3, e0), e3)) & (e3 = op(op(e3, e0), e0)) & ~ (op(e0, e3) = op(e3, e0))) | ~ sP15), inference(nnf_transformation, [], [f262])).
fof(f1441, plain, (~ spl24_160 | spl24_162), inference(avatar_split_clause, [], [f517, f1438, f1429])).
fof(f517, plain, ((e3 = op(op(e3, e0), e0)) | ~ sP15), inference(cnf_transformation, [], [f280])).
fof(f1436, plain, (~ spl24_160 | ~ spl24_161), inference(avatar_split_clause, [], [f518, f1433, f1429])).
fof(f518, plain, (~ (e0 = op(op(e3, e0), e3)) | ~ sP15), inference(cnf_transformation, [], [f280])).
fof(f1427, plain, (~ spl24_156 | ~ spl24_159), inference(avatar_split_clause, [], [f513, f1424, f1410])).
fof(f513, plain, (~ (op(e1, e3) = op(e3, e1)) | ~ sP16), inference(cnf_transformation, [], [f279])).
fof(f279, plain, ((~ (e1 = op(op(e3, e1), e3)) & (e3 = op(op(e3, e1), e1)) & ~ (op(e1, e3) = op(e3, e1))) | ~ sP16), inference(nnf_transformation, [], [f263])).
fof(f1422, plain, (~ spl24_156 | spl24_158), inference(avatar_split_clause, [], [f514, f1419, f1410])).
fof(f514, plain, ((e3 = op(op(e3, e1), e1)) | ~ sP16), inference(cnf_transformation, [], [f279])).
fof(f1417, plain, (~ spl24_156 | ~ spl24_157), inference(avatar_split_clause, [], [f515, f1414, f1410])).
fof(f515, plain, (~ (e1 = op(op(e3, e1), e3)) | ~ sP16), inference(cnf_transformation, [], [f279])).
fof(f1408, plain, (~ spl24_152 | ~ spl24_155), inference(avatar_split_clause, [], [f510, f1405, f1391])).
fof(f510, plain, (~ (op(e2, e3) = op(e3, e2)) | ~ sP17), inference(cnf_transformation, [], [f278])).
fof(f278, plain, ((~ (e2 = op(op(e3, e2), e3)) & (e3 = op(op(e3, e2), e2)) & ~ (op(e2, e3) = op(e3, e2))) | ~ sP17), inference(nnf_transformation, [], [f264])).
fof(f1403, plain, (~ spl24_152 | spl24_154), inference(avatar_split_clause, [], [f511, f1400, f1391])).
fof(f511, plain, ((e3 = op(op(e3, e2), e2)) | ~ sP17), inference(cnf_transformation, [], [f278])).
fof(f1398, plain, (~ spl24_152 | ~ spl24_153), inference(avatar_split_clause, [], [f512, f1395, f1391])).
fof(f512, plain, (~ (e2 = op(op(e3, e2), e3)) | ~ sP17), inference(cnf_transformation, [], [f278])).
fof(f1389, plain, ~ spl24_150, inference(avatar_split_clause, [], [f691, f1380])).
fof(f691, plain, ~ sP18, inference(trivial_inequality_removal, [], [f507])).
fof(f507, plain, (~ (op(e3, e3) = op(e3, e3)) | ~ sP18), inference(cnf_transformation, [], [f277])).
fof(f277, plain, ((~ (e3 = op(op(e3, e3), e3)) & (e3 = op(op(e3, e3), e3)) & ~ (op(e3, e3) = op(e3, e3))) | ~ sP18), inference(nnf_transformation, [], [f265])).
fof(f1378, plain, (~ spl24_147 | ~ spl24_134), inference(avatar_split_clause, [], [f504, f1303, f1365])).
fof(f504, plain, (~ (op(e3, e4) = op(e4, e3)) | ~ sP19), inference(cnf_transformation, [], [f276])).
fof(f276, plain, ((~ (e4 = op(op(e3, e4), e3)) & (e3 = op(op(e3, e4), e4)) & ~ (op(e3, e4) = op(e4, e3))) | ~ sP19), inference(nnf_transformation, [], [f266])).
fof(f1377, plain, (~ spl24_147 | spl24_149), inference(avatar_split_clause, [], [f505, f1374, f1365])).
fof(f505, plain, ((e3 = op(op(e3, e4), e4)) | ~ sP19), inference(cnf_transformation, [], [f276])).
fof(f1372, plain, (~ spl24_147 | ~ spl24_148), inference(avatar_split_clause, [], [f506, f1369, f1365])).
fof(f506, plain, (~ (e4 = op(op(e3, e4), e3)) | ~ sP19), inference(cnf_transformation, [], [f276])).
fof(f1363, plain, (~ spl24_143 | ~ spl24_146), inference(avatar_split_clause, [], [f501, f1360, f1346])).
fof(f501, plain, (~ (op(e0, e4) = op(e4, e0)) | ~ sP20), inference(cnf_transformation, [], [f275])).
fof(f275, plain, ((~ (e0 = op(op(e4, e0), e4)) & (e4 = op(op(e4, e0), e0)) & ~ (op(e0, e4) = op(e4, e0))) | ~ sP20), inference(nnf_transformation, [], [f267])).
fof(f1358, plain, (~ spl24_143 | spl24_145), inference(avatar_split_clause, [], [f502, f1355, f1346])).
fof(f502, plain, ((e4 = op(op(e4, e0), e0)) | ~ sP20), inference(cnf_transformation, [], [f275])).
fof(f1353, plain, (~ spl24_143 | ~ spl24_144), inference(avatar_split_clause, [], [f503, f1350, f1346])).
fof(f503, plain, (~ (e0 = op(op(e4, e0), e4)) | ~ sP20), inference(cnf_transformation, [], [f275])).
fof(f1344, plain, (~ spl24_139 | ~ spl24_142), inference(avatar_split_clause, [], [f498, f1341, f1327])).
fof(f498, plain, (~ (op(e1, e4) = op(e4, e1)) | ~ sP21), inference(cnf_transformation, [], [f274])).
fof(f274, plain, ((~ (e1 = op(op(e4, e1), e4)) & (e4 = op(op(e4, e1), e1)) & ~ (op(e1, e4) = op(e4, e1))) | ~ sP21), inference(nnf_transformation, [], [f268])).
fof(f1339, plain, (~ spl24_139 | spl24_141), inference(avatar_split_clause, [], [f499, f1336, f1327])).
fof(f499, plain, ((e4 = op(op(e4, e1), e1)) | ~ sP21), inference(cnf_transformation, [], [f274])).
fof(f1334, plain, (~ spl24_139 | ~ spl24_140), inference(avatar_split_clause, [], [f500, f1331, f1327])).
fof(f500, plain, (~ (e1 = op(op(e4, e1), e4)) | ~ sP21), inference(cnf_transformation, [], [f274])).
fof(f1325, plain, (~ spl24_135 | ~ spl24_138), inference(avatar_split_clause, [], [f495, f1322, f1308])).
fof(f495, plain, (~ (op(e2, e4) = op(e4, e2)) | ~ sP22), inference(cnf_transformation, [], [f273])).
fof(f273, plain, ((~ (e2 = op(op(e4, e2), e4)) & (e4 = op(op(e4, e2), e2)) & ~ (op(e2, e4) = op(e4, e2))) | ~ sP22), inference(nnf_transformation, [], [f269])).
fof(f1320, plain, (~ spl24_135 | spl24_137), inference(avatar_split_clause, [], [f496, f1317, f1308])).
fof(f496, plain, ((e4 = op(op(e4, e2), e2)) | ~ sP22), inference(cnf_transformation, [], [f273])).
fof(f1315, plain, (~ spl24_135 | ~ spl24_136), inference(avatar_split_clause, [], [f497, f1312, f1308])).
fof(f497, plain, (~ (e2 = op(op(e4, e2), e4)) | ~ sP22), inference(cnf_transformation, [], [f273])).
fof(f1306, plain, (~ spl24_131 | ~ spl24_134), inference(avatar_split_clause, [], [f492, f1303, f1289])).
fof(f492, plain, (~ (op(e3, e4) = op(e4, e3)) | ~ sP23), inference(cnf_transformation, [], [f272])).
fof(f272, plain, ((~ (e3 = op(op(e4, e3), e4)) & (e4 = op(op(e4, e3), e3)) & ~ (op(e3, e4) = op(e4, e3))) | ~ sP23), inference(nnf_transformation, [], [f270])).
fof(f1301, plain, (~ spl24_131 | spl24_133), inference(avatar_split_clause, [], [f493, f1298, f1289])).
fof(f493, plain, ((e4 = op(op(e4, e3), e3)) | ~ sP23), inference(cnf_transformation, [], [f272])).
fof(f1296, plain, (~ spl24_131 | ~ spl24_132), inference(avatar_split_clause, [], [f494, f1293, f1289])).
fof(f494, plain, (~ (e3 = op(op(e4, e3), e4)) | ~ sP23), inference(cnf_transformation, [], [f272])).
fof(f1285, plain, (spl24_122 | spl24_117 | spl24_112 | spl24_107 | spl24_102), inference(avatar_split_clause, [], [f334, f1117, f1138, f1159, f1180, f1201])).
fof(f334, plain, ((e1 = op(e0, e4)) | (e1 = op(e0, e3)) | (e1 = op(e0, e2)) | (e1 = op(e0, e1)) | (op(e0, e0) = e1)), inference(cnf_transformation, [], [f3])).
fof(f3, plain, (((e4 = op(e4, e4)) | (e4 = op(e3, e4)) | (e4 = op(e2, e4)) | (e4 = op(e1, e4)) | (e4 = op(e0, e4))) & ((e4 = op(e4, e4)) | (e4 = op(e4, e3)) | (e4 = op(e4, e2)) | (e4 = op(e4, e1)) | (e4 = op(e4, e0))) & ((e3 = op(e4, e4)) | (e3 = op(e3, e4)) | (e3 = op(e2, e4)) | (e3 = op(e1, e4)) | (e3 = op(e0, e4))) & ((e3 = op(e4, e4)) | (e3 = op(e4, e3)) | (e3 = op(e4, e2)) | (e3 = op(e4, e1)) | (e3 = op(e4, e0))) & ((e2 = op(e4, e4)) | (e2 = op(e3, e4)) | (e2 = op(e2, e4)) | (e2 = op(e1, e4)) | (e2 = op(e0, e4))) & ((e2 = op(e4, e4)) | (e2 = op(e4, e3)) | (e2 = op(e4, e2)) | (e2 = op(e4, e1)) | (e2 = op(e4, e0))) & ((e1 = op(e4, e4)) | (e1 = op(e3, e4)) | (e1 = op(e2, e4)) | (e1 = op(e1, e4)) | (e1 = op(e0, e4))) & ((e1 = op(e4, e4)) | (e1 = op(e4, e3)) | (e1 = op(e4, e2)) | (e1 = op(e4, e1)) | (e1 = op(e4, e0))) & ((e0 = op(e4, e4)) | (e0 = op(e3, e4)) | (e0 = op(e2, e4)) | (e0 = op(e1, e4)) | (e0 = op(e0, e4))) & ((e0 = op(e4, e4)) | (e0 = op(e4, e3)) | (e0 = op(e4, e2)) | (e0 = op(e4, e1)) | (e0 = op(e4, e0))) & ((e4 = op(e4, e3)) | (e4 = op(e3, e3)) | (e4 = op(e2, e3)) | (e4 = op(e1, e3)) | (e4 = op(e0, e3))) & ((e4 = op(e3, e4)) | (e4 = op(e3, e3)) | (e4 = op(e3, e2)) | (e4 = op(e3, e1)) | (e4 = op(e3, e0))) & ((e3 = op(e4, e3)) | (e3 = op(e3, e3)) | (e3 = op(e2, e3)) | (e3 = op(e1, e3)) | (e3 = op(e0, e3))) & ((e3 = op(e3, e4)) | (e3 = op(e3, e3)) | (e3 = op(e3, e2)) | (e3 = op(e3, e1)) | (e3 = op(e3, e0))) & ((e2 = op(e4, e3)) | (e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(e0, e3))) & ((e2 = op(e3, e4)) | (e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, e0))) & ((e1 = op(e4, e3)) | (e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(e0, e3))) & ((e1 = op(e3, e4)) | (e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, e0))) & ((e0 = op(e4, e3)) | (e0 = op(e3, e3)) | (e0 = op(e2, e3)) | (e0 = op(e1, e3)) | (e0 = op(e0, e3))) & ((e0 = op(e3, e4)) | (e0 = op(e3, e3)) | (e0 = op(e3, e2)) | (e0 = op(e3, e1)) | (e0 = op(e3, e0))) & ((e4 = op(e4, e2)) | (e4 = op(e3, e2)) | (e4 = op(e2, e2)) | (e4 = op(e1, e2)) | (e4 = op(e0, e2))) & ((e4 = op(e2, e4)) | (e4 = op(e2, e3)) | (e4 = op(e2, e2)) | (e4 = op(e2, e1)) | (e4 = op(e2, e0))) & ((e3 = op(e4, e2)) | (e3 = op(e3, e2)) | (e3 = op(e2, e2)) | (e3 = op(e1, e2)) | (e3 = op(e0, e2))) & ((e3 = op(e2, e4)) | (e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, e0))) & ((e2 = op(e4, e2)) | (e2 = op(e3, e2)) | (e2 = op(e2, e2)) | (e2 = op(e1, e2)) | (e2 = op(e0, e2))) & ((e2 = op(e2, e4)) | (e2 = op(e2, e3)) | (e2 = op(e2, e2)) | (e2 = op(e2, e1)) | (e2 = op(e2, e0))) & ((e1 = op(e4, e2)) | (e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(e0, e2))) & ((e1 = op(e2, e4)) | (e1 = op(e2, e3)) | (e1 = op(e2, e2)) | (e1 = op(e2, e1)) | (e1 = op(e2, e0))) & ((e0 = op(e4, e2)) | (e0 = op(e3, e2)) | (e0 = op(e2, e2)) | (e0 = op(e1, e2)) | (e0 = op(e0, e2))) & ((e0 = op(e2, e4)) | (e0 = op(e2, e3)) | (e0 = op(e2, e2)) | (e0 = op(e2, e1)) | (e0 = op(e2, e0))) & ((e4 = op(e4, e1)) | (e4 = op(e3, e1)) | (e4 = op(e2, e1)) | (e4 = op(e1, e1)) | (e4 = op(e0, e1))) & ((e4 = op(e1, e4)) | (e4 = op(e1, e3)) | (e4 = op(e1, e2)) | (e4 = op(e1, e1)) | (e4 = op(e1, e0))) & ((e3 = op(e4, e1)) | (e3 = op(e3, e1)) | (e3 = op(e2, e1)) | (e3 = op(e1, e1)) | (e3 = op(e0, e1))) & ((e3 = op(e1, e4)) | (e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, e0))) & ((e2 = op(e4, e1)) | (e2 = op(e3, e1)) | (e2 = op(e2, e1)) | (e2 = op(e1, e1)) | (e2 = op(e0, e1))) & ((e2 = op(e1, e4)) | (e2 = op(e1, e3)) | (e2 = op(e1, e2)) | (e2 = op(e1, e1)) | (e2 = op(e1, e0))) & ((e1 = op(e4, e1)) | (e1 = op(e3, e1)) | (e1 = op(e2, e1)) | (e1 = op(e1, e1)) | (e1 = op(e0, e1))) & ((e1 = op(e1, e4)) | (e1 = op(e1, e3)) | (e1 = op(e1, e2)) | (e1 = op(e1, e1)) | (e1 = op(e1, e0))) & ((e0 = op(e4, e1)) | (e0 = op(e3, e1)) | (e0 = op(e2, e1)) | (e0 = op(e1, e1)) | (e0 = op(e0, e1))) & ((e0 = op(e1, e4)) | (e0 = op(e1, e3)) | (e0 = op(e1, e2)) | (e0 = op(e1, e1)) | (e0 = op(e1, e0))) & ((e4 = op(e4, e0)) | (e4 = op(e3, e0)) | (e4 = op(e2, e0)) | (e4 = op(e1, e0)) | (op(e0, e0) = e4)) & ((e4 = op(e0, e4)) | (e4 = op(e0, e3)) | (e4 = op(e0, e2)) | (e4 = op(e0, e1)) | (op(e0, e0) = e4)) & ((e3 = op(e4, e0)) | (e3 = op(e3, e0)) | (e3 = op(e2, e0)) | (e3 = op(e1, e0)) | (op(e0, e0) = e3)) & ((e3 = op(e0, e4)) | (e3 = op(e0, e3)) | (e3 = op(e0, e2)) | (e3 = op(e0, e1)) | (op(e0, e0) = e3)) & ((e2 = op(e4, e0)) | (e2 = op(e3, e0)) | (e2 = op(e2, e0)) | (e2 = op(e1, e0)) | (op(e0, e0) = e2)) & ((e2 = op(e0, e4)) | (e2 = op(e0, e3)) | (e2 = op(e0, e2)) | (e2 = op(e0, e1)) | (op(e0, e0) = e2)) & ((e1 = op(e4, e0)) | (e1 = op(e3, e0)) | (e1 = op(e2, e0)) | (e1 = op(e1, e0)) | (op(e0, e0) = e1)) & ((e1 = op(e0, e4)) | (e1 = op(e0, e3)) | (e1 = op(e0, e2)) | (e1 = op(e0, e1)) | (op(e0, e0) = e1)) & ((e0 = op(e4, e0)) | (e0 = op(e3, e0)) | (e0 = op(e2, e0)) | (e0 = op(e1, e0)) | (e0 = op(e0, e0))) & ((e0 = op(e0, e4)) | (e0 = op(e0, e3)) | (e0 = op(e0, e2)) | (e0 = op(e0, e1)) | (e0 = op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax3)).
fof(f1284, plain, (spl24_122 | spl24_97 | spl24_72 | spl24_47 | spl24_22), inference(avatar_split_clause, [], [f335, f781, f886, f991, f1096, f1201])).
fof(f335, plain, ((e1 = op(e4, e0)) | (e1 = op(e3, e0)) | (e1 = op(e2, e0)) | (e1 = op(e1, e0)) | (op(e0, e0) = e1)), inference(cnf_transformation, [], [f3])).
fof(f1283, plain, (spl24_123 | spl24_118 | spl24_113 | spl24_108 | spl24_103), inference(avatar_split_clause, [], [f336, f1121, f1142, f1163, f1184, f1205])).
fof(f336, plain, ((e2 = op(e0, e4)) | (e2 = op(e0, e3)) | (e2 = op(e0, e2)) | (e2 = op(e0, e1)) | (op(e0, e0) = e2)), inference(cnf_transformation, [], [f3])).
fof(f1282, plain, (spl24_123 | spl24_98 | spl24_73 | spl24_48 | spl24_23), inference(avatar_split_clause, [], [f337, f785, f890, f995, f1100, f1205])).
fof(f337, plain, ((e2 = op(e4, e0)) | (e2 = op(e3, e0)) | (e2 = op(e2, e0)) | (e2 = op(e1, e0)) | (op(e0, e0) = e2)), inference(cnf_transformation, [], [f3])).
fof(f1281, plain, (spl24_124 | spl24_119 | spl24_114 | spl24_109 | spl24_104), inference(avatar_split_clause, [], [f338, f1125, f1146, f1167, f1188, f1209])).
fof(f338, plain, ((e3 = op(e0, e4)) | (e3 = op(e0, e3)) | (e3 = op(e0, e2)) | (e3 = op(e0, e1)) | (op(e0, e0) = e3)), inference(cnf_transformation, [], [f3])).
fof(f1280, plain, (spl24_124 | spl24_99 | spl24_74 | spl24_49 | spl24_24), inference(avatar_split_clause, [], [f339, f789, f894, f999, f1104, f1209])).
fof(f339, plain, ((e3 = op(e4, e0)) | (e3 = op(e3, e0)) | (e3 = op(e2, e0)) | (e3 = op(e1, e0)) | (op(e0, e0) = e3)), inference(cnf_transformation, [], [f3])).
fof(f1279, plain, (spl24_125 | spl24_120 | spl24_115 | spl24_110 | spl24_105), inference(avatar_split_clause, [], [f340, f1129, f1150, f1171, f1192, f1213])).
fof(f340, plain, ((e4 = op(e0, e4)) | (e4 = op(e0, e3)) | (e4 = op(e0, e2)) | (e4 = op(e0, e1)) | (op(e0, e0) = e4)), inference(cnf_transformation, [], [f3])).
fof(f1278, plain, (spl24_125 | spl24_100 | spl24_75 | spl24_50 | spl24_25), inference(avatar_split_clause, [], [f341, f793, f898, f1003, f1108, f1213])).
fof(f341, plain, ((e4 = op(e4, e0)) | (e4 = op(e3, e0)) | (e4 = op(e2, e0)) | (e4 = op(e1, e0)) | (op(e0, e0) = e4)), inference(cnf_transformation, [], [f3])).
fof(f1277, plain, (spl24_96 | spl24_91 | spl24_86 | spl24_81 | spl24_76), inference(avatar_split_clause, [], [f342, f1008, f1029, f1050, f1071, f1092])).
fof(f342, plain, ((e0 = op(e1, e4)) | (e0 = op(e1, e3)) | (e0 = op(e1, e2)) | (e0 = op(e1, e1)) | (e0 = op(e1, e0))), inference(cnf_transformation, [], [f3])).
fof(f1276, plain, (spl24_116 | spl24_91 | spl24_66 | spl24_41 | spl24_16), inference(avatar_split_clause, [], [f343, f756, f861, f966, f1071, f1176])).
fof(f343, plain, ((e0 = op(e4, e1)) | (e0 = op(e3, e1)) | (e0 = op(e2, e1)) | (e0 = op(e1, e1)) | (e0 = op(e0, e1))), inference(cnf_transformation, [], [f3])).
fof(f1273, plain, (spl24_98 | spl24_93 | spl24_88 | spl24_83 | spl24_78), inference(avatar_split_clause, [], [f346, f1016, f1037, f1058, f1079, f1100])).
fof(f346, plain, ((e2 = op(e1, e4)) | (e2 = op(e1, e3)) | (e2 = op(e1, e2)) | (e2 = op(e1, e1)) | (e2 = op(e1, e0))), inference(cnf_transformation, [], [f3])).
fof(f1272, plain, (spl24_118 | spl24_93 | spl24_68 | spl24_43 | spl24_18), inference(avatar_split_clause, [], [f347, f764, f869, f974, f1079, f1184])).
fof(f347, plain, ((e2 = op(e4, e1)) | (e2 = op(e3, e1)) | (e2 = op(e2, e1)) | (e2 = op(e1, e1)) | (e2 = op(e0, e1))), inference(cnf_transformation, [], [f3])).
fof(f1271, plain, (spl24_99 | spl24_94 | spl24_89 | spl24_84 | spl24_79), inference(avatar_split_clause, [], [f348, f1020, f1041, f1062, f1083, f1104])).
fof(f348, plain, ((e3 = op(e1, e4)) | (e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, e0))), inference(cnf_transformation, [], [f3])).
fof(f1270, plain, (spl24_119 | spl24_94 | spl24_69 | spl24_44 | spl24_19), inference(avatar_split_clause, [], [f349, f768, f873, f978, f1083, f1188])).
fof(f349, plain, ((e3 = op(e4, e1)) | (e3 = op(e3, e1)) | (e3 = op(e2, e1)) | (e3 = op(e1, e1)) | (e3 = op(e0, e1))), inference(cnf_transformation, [], [f3])).
fof(f1269, plain, (spl24_100 | spl24_95 | spl24_90 | spl24_85 | spl24_80), inference(avatar_split_clause, [], [f350, f1024, f1045, f1066, f1087, f1108])).
fof(f350, plain, ((e4 = op(e1, e4)) | (e4 = op(e1, e3)) | (e4 = op(e1, e2)) | (e4 = op(e1, e1)) | (e4 = op(e1, e0))), inference(cnf_transformation, [], [f3])).
fof(f1268, plain, (spl24_120 | spl24_95 | spl24_70 | spl24_45 | spl24_20), inference(avatar_split_clause, [], [f351, f772, f877, f982, f1087, f1192])).
fof(f351, plain, ((e4 = op(e4, e1)) | (e4 = op(e3, e1)) | (e4 = op(e2, e1)) | (e4 = op(e1, e1)) | (e4 = op(e0, e1))), inference(cnf_transformation, [], [f3])).
fof(f1267, plain, (spl24_71 | spl24_66 | spl24_61 | spl24_56 | spl24_51), inference(avatar_split_clause, [], [f352, f903, f924, f945, f966, f987])).
fof(f352, plain, ((e0 = op(e2, e4)) | (e0 = op(e2, e3)) | (e0 = op(e2, e2)) | (e0 = op(e2, e1)) | (e0 = op(e2, e0))), inference(cnf_transformation, [], [f3])).
fof(f1266, plain, (spl24_111 | spl24_86 | spl24_61 | spl24_36 | spl24_11), inference(avatar_split_clause, [], [f353, f735, f840, f945, f1050, f1155])).
fof(f353, plain, ((e0 = op(e4, e2)) | (e0 = op(e3, e2)) | (e0 = op(e2, e2)) | (e0 = op(e1, e2)) | (e0 = op(e0, e2))), inference(cnf_transformation, [], [f3])).
fof(f1265, plain, (spl24_72 | spl24_67 | spl24_62 | spl24_57 | spl24_52), inference(avatar_split_clause, [], [f354, f907, f928, f949, f970, f991])).
fof(f354, plain, ((e1 = op(e2, e4)) | (e1 = op(e2, e3)) | (e1 = op(e2, e2)) | (e1 = op(e2, e1)) | (e1 = op(e2, e0))), inference(cnf_transformation, [], [f3])).
fof(f1264, plain, (spl24_112 | spl24_87 | spl24_62 | spl24_37 | spl24_12), inference(avatar_split_clause, [], [f355, f739, f844, f949, f1054, f1159])).
fof(f355, plain, ((e1 = op(e4, e2)) | (e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(e0, e2))), inference(cnf_transformation, [], [f3])).
fof(f1261, plain, (spl24_74 | spl24_69 | spl24_64 | spl24_59 | spl24_54), inference(avatar_split_clause, [], [f358, f915, f936, f957, f978, f999])).
fof(f358, plain, ((e3 = op(e2, e4)) | (e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, e0))), inference(cnf_transformation, [], [f3])).
fof(f1260, plain, (spl24_114 | spl24_89 | spl24_64 | spl24_39 | spl24_14), inference(avatar_split_clause, [], [f359, f747, f852, f957, f1062, f1167])).
fof(f359, plain, ((e3 = op(e4, e2)) | (e3 = op(e3, e2)) | (e3 = op(e2, e2)) | (e3 = op(e1, e2)) | (e3 = op(e0, e2))), inference(cnf_transformation, [], [f3])).
fof(f1259, plain, (spl24_75 | spl24_70 | spl24_65 | spl24_60 | spl24_55), inference(avatar_split_clause, [], [f360, f919, f940, f961, f982, f1003])).
fof(f360, plain, ((e4 = op(e2, e4)) | (e4 = op(e2, e3)) | (e4 = op(e2, e2)) | (e4 = op(e2, e1)) | (e4 = op(e2, e0))), inference(cnf_transformation, [], [f3])).
fof(f1258, plain, (spl24_115 | spl24_90 | spl24_65 | spl24_40 | spl24_15), inference(avatar_split_clause, [], [f361, f751, f856, f961, f1066, f1171])).
fof(f361, plain, ((e4 = op(e4, e2)) | (e4 = op(e3, e2)) | (e4 = op(e2, e2)) | (e4 = op(e1, e2)) | (e4 = op(e0, e2))), inference(cnf_transformation, [], [f3])).
fof(f1257, plain, (spl24_46 | spl24_41 | spl24_36 | spl24_31 | spl24_26), inference(avatar_split_clause, [], [f362, f798, f819, f840, f861, f882])).
fof(f362, plain, ((e0 = op(e3, e4)) | (e0 = op(e3, e3)) | (e0 = op(e3, e2)) | (e0 = op(e3, e1)) | (e0 = op(e3, e0))), inference(cnf_transformation, [], [f3])).
fof(f1256, plain, (spl24_106 | spl24_81 | spl24_56 | spl24_31 | spl24_6), inference(avatar_split_clause, [], [f363, f714, f819, f924, f1029, f1134])).
fof(f363, plain, ((e0 = op(e4, e3)) | (e0 = op(e3, e3)) | (e0 = op(e2, e3)) | (e0 = op(e1, e3)) | (e0 = op(e0, e3))), inference(cnf_transformation, [], [f3])).
fof(f1255, plain, (spl24_47 | spl24_42 | spl24_37 | spl24_32 | spl24_27), inference(avatar_split_clause, [], [f364, f802, f823, f844, f865, f886])).
fof(f364, plain, ((e1 = op(e3, e4)) | (e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, e0))), inference(cnf_transformation, [], [f3])).
fof(f1254, plain, (spl24_107 | spl24_82 | spl24_57 | spl24_32 | spl24_7), inference(avatar_split_clause, [], [f365, f718, f823, f928, f1033, f1138])).
fof(f365, plain, ((e1 = op(e4, e3)) | (e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(e0, e3))), inference(cnf_transformation, [], [f3])).
fof(f1253, plain, (spl24_48 | spl24_43 | spl24_38 | spl24_33 | spl24_28), inference(avatar_split_clause, [], [f366, f806, f827, f848, f869, f890])).
fof(f366, plain, ((e2 = op(e3, e4)) | (e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, e0))), inference(cnf_transformation, [], [f3])).
fof(f1252, plain, (spl24_108 | spl24_83 | spl24_58 | spl24_33 | spl24_8), inference(avatar_split_clause, [], [f367, f722, f827, f932, f1037, f1142])).
fof(f367, plain, ((e2 = op(e4, e3)) | (e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(e0, e3))), inference(cnf_transformation, [], [f3])).
fof(f1249, plain, (spl24_50 | spl24_45 | spl24_40 | spl24_35 | spl24_30), inference(avatar_split_clause, [], [f370, f814, f835, f856, f877, f898])).
fof(f370, plain, ((e4 = op(e3, e4)) | (e4 = op(e3, e3)) | (e4 = op(e3, e2)) | (e4 = op(e3, e1)) | (e4 = op(e3, e0))), inference(cnf_transformation, [], [f3])).
fof(f1248, plain, (spl24_110 | spl24_85 | spl24_60 | spl24_35 | spl24_10), inference(avatar_split_clause, [], [f371, f730, f835, f940, f1045, f1150])).
fof(f371, plain, ((e4 = op(e4, e3)) | (e4 = op(e3, e3)) | (e4 = op(e2, e3)) | (e4 = op(e1, e3)) | (e4 = op(e0, e3))), inference(cnf_transformation, [], [f3])).
fof(f1247, plain, (spl24_21 | spl24_16 | spl24_11 | spl24_6 | spl24_1), inference(avatar_split_clause, [], [f372, f693, f714, f735, f756, f777])).
fof(f372, plain, ((e0 = op(e4, e4)) | (e0 = op(e4, e3)) | (e0 = op(e4, e2)) | (e0 = op(e4, e1)) | (e0 = op(e4, e0))), inference(cnf_transformation, [], [f3])).
fof(f1246, plain, (spl24_101 | spl24_76 | spl24_51 | spl24_26 | spl24_1), inference(avatar_split_clause, [], [f373, f693, f798, f903, f1008, f1113])).
fof(f373, plain, ((e0 = op(e4, e4)) | (e0 = op(e3, e4)) | (e0 = op(e2, e4)) | (e0 = op(e1, e4)) | (e0 = op(e0, e4))), inference(cnf_transformation, [], [f3])).
fof(f1245, plain, (spl24_22 | spl24_17 | spl24_12 | spl24_7 | spl24_2), inference(avatar_split_clause, [], [f374, f697, f718, f739, f760, f781])).
fof(f374, plain, ((e1 = op(e4, e4)) | (e1 = op(e4, e3)) | (e1 = op(e4, e2)) | (e1 = op(e4, e1)) | (e1 = op(e4, e0))), inference(cnf_transformation, [], [f3])).
fof(f1244, plain, (spl24_102 | spl24_77 | spl24_52 | spl24_27 | spl24_2), inference(avatar_split_clause, [], [f375, f697, f802, f907, f1012, f1117])).
fof(f375, plain, ((e1 = op(e4, e4)) | (e1 = op(e3, e4)) | (e1 = op(e2, e4)) | (e1 = op(e1, e4)) | (e1 = op(e0, e4))), inference(cnf_transformation, [], [f3])).
fof(f1243, plain, (spl24_23 | spl24_18 | spl24_13 | spl24_8 | spl24_3), inference(avatar_split_clause, [], [f376, f701, f722, f743, f764, f785])).
fof(f376, plain, ((e2 = op(e4, e4)) | (e2 = op(e4, e3)) | (e2 = op(e4, e2)) | (e2 = op(e4, e1)) | (e2 = op(e4, e0))), inference(cnf_transformation, [], [f3])).
fof(f1242, plain, (spl24_103 | spl24_78 | spl24_53 | spl24_28 | spl24_3), inference(avatar_split_clause, [], [f377, f701, f806, f911, f1016, f1121])).
fof(f377, plain, ((e2 = op(e4, e4)) | (e2 = op(e3, e4)) | (e2 = op(e2, e4)) | (e2 = op(e1, e4)) | (e2 = op(e0, e4))), inference(cnf_transformation, [], [f3])).
fof(f1241, plain, (spl24_24 | spl24_19 | spl24_14 | spl24_9 | spl24_4), inference(avatar_split_clause, [], [f378, f705, f726, f747, f768, f789])).
fof(f378, plain, ((e3 = op(e4, e4)) | (e3 = op(e4, e3)) | (e3 = op(e4, e2)) | (e3 = op(e4, e1)) | (e3 = op(e4, e0))), inference(cnf_transformation, [], [f3])).
fof(f1240, plain, (spl24_104 | spl24_79 | spl24_54 | spl24_29 | spl24_4), inference(avatar_split_clause, [], [f379, f705, f810, f915, f1020, f1125])).
fof(f379, plain, ((e3 = op(e4, e4)) | (e3 = op(e3, e4)) | (e3 = op(e2, e4)) | (e3 = op(e1, e4)) | (e3 = op(e0, e4))), inference(cnf_transformation, [], [f3])).
fof(f1237, plain, (spl24_126 | spl24_127 | spl24_128 | spl24_129 | spl24_130), inference(avatar_split_clause, [], [f331, f1234, f1230, f1226, f1222, f1218])).
fof(f331, plain, ((e4 = unit) | (e3 = unit) | (e2 = unit) | (e1 = unit) | (e0 = unit)), inference(cnf_transformation, [], [f2])).
fof(f1216, plain, (spl24_121 | spl24_122 | spl24_123 | spl24_124 | spl24_125), inference(avatar_split_clause, [], [f296, f1213, f1209, f1205, f1201, f1197])).
fof(f296, plain, ((op(e0, e0) = e4) | (op(e0, e0) = e3) | (op(e0, e0) = e2) | (op(e0, e0) = e1) | (e0 = op(e0, e0))), inference(cnf_transformation, [], [f1])).
fof(f1, plain, (((e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | (e1 = op(e4, e4)) | (e0 = op(e4, e4))) & ((e4 = op(e4, e3)) | (e3 = op(e4, e3)) | (e2 = op(e4, e3)) | (e1 = op(e4, e3)) | (e0 = op(e4, e3))) & ((e4 = op(e4, e2)) | (e3 = op(e4, e2)) | (e2 = op(e4, e2)) | (e1 = op(e4, e2)) | (e0 = op(e4, e2))) & ((e4 = op(e4, e1)) | (e3 = op(e4, e1)) | (e2 = op(e4, e1)) | (e1 = op(e4, e1)) | (e0 = op(e4, e1))) & ((e4 = op(e4, e0)) | (e3 = op(e4, e0)) | (e2 = op(e4, e0)) | (e1 = op(e4, e0)) | (e0 = op(e4, e0))) & ((e4 = op(e3, e4)) | (e3 = op(e3, e4)) | (e2 = op(e3, e4)) | (e1 = op(e3, e4)) | (e0 = op(e3, e4))) & ((e4 = op(e3, e3)) | (e3 = op(e3, e3)) | (e2 = op(e3, e3)) | (e1 = op(e3, e3)) | (e0 = op(e3, e3))) & ((e4 = op(e3, e2)) | (e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (e0 = op(e3, e2))) & ((e4 = op(e3, e1)) | (e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (e0 = op(e3, e1))) & ((e4 = op(e3, e0)) | (e3 = op(e3, e0)) | (e2 = op(e3, e0)) | (e1 = op(e3, e0)) | (e0 = op(e3, e0))) & ((e4 = op(e2, e4)) | (e3 = op(e2, e4)) | (e2 = op(e2, e4)) | (e1 = op(e2, e4)) | (e0 = op(e2, e4))) & ((e4 = op(e2, e3)) | (e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))) & ((e4 = op(e2, e2)) | (e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (e0 = op(e2, e2))) & ((e4 = op(e2, e1)) | (e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))) & ((e4 = op(e2, e0)) | (e3 = op(e2, e0)) | (e2 = op(e2, e0)) | (e1 = op(e2, e0)) | (e0 = op(e2, e0))) & ((e4 = op(e1, e4)) | (e3 = op(e1, e4)) | (e2 = op(e1, e4)) | (e1 = op(e1, e4)) | (e0 = op(e1, e4))) & ((e4 = op(e1, e3)) | (e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (e0 = op(e1, e3))) & ((e4 = op(e1, e2)) | (e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (e0 = op(e1, e2))) & ((e4 = op(e1, e1)) | (e3 = op(e1, e1)) | (e2 = op(e1, e1)) | (e1 = op(e1, e1)) | (e0 = op(e1, e1))) & ((e4 = op(e1, e0)) | (e3 = op(e1, e0)) | (e2 = op(e1, e0)) | (e1 = op(e1, e0)) | (e0 = op(e1, e0))) & ((e4 = op(e0, e4)) | (e3 = op(e0, e4)) | (e2 = op(e0, e4)) | (e1 = op(e0, e4)) | (e0 = op(e0, e4))) & ((e4 = op(e0, e3)) | (e3 = op(e0, e3)) | (e2 = op(e0, e3)) | (e1 = op(e0, e3)) | (e0 = op(e0, e3))) & ((e4 = op(e0, e2)) | (e3 = op(e0, e2)) | (e2 = op(e0, e2)) | (e1 = op(e0, e2)) | (e0 = op(e0, e2))) & ((e4 = op(e0, e1)) | (e3 = op(e0, e1)) | (e2 = op(e0, e1)) | (e1 = op(e0, e1)) | (e0 = op(e0, e1))) & ((op(e0, e0) = e4) | (op(e0, e0) = e3) | (op(e0, e0) = e2) | (op(e0, e0) = e1) | (e0 = op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG101+1.p', ax1)).
fof(f1195, plain, (spl24_116 | spl24_117 | spl24_118 | spl24_119 | spl24_120), inference(avatar_split_clause, [], [f297, f1192, f1188, f1184, f1180, f1176])).
fof(f297, plain, ((e4 = op(e0, e1)) | (e3 = op(e0, e1)) | (e2 = op(e0, e1)) | (e1 = op(e0, e1)) | (e0 = op(e0, e1))), inference(cnf_transformation, [], [f1])).
fof(f1174, plain, (spl24_111 | spl24_112 | spl24_113 | spl24_114 | spl24_115), inference(avatar_split_clause, [], [f298, f1171, f1167, f1163, f1159, f1155])).
fof(f298, plain, ((e4 = op(e0, e2)) | (e3 = op(e0, e2)) | (e2 = op(e0, e2)) | (e1 = op(e0, e2)) | (e0 = op(e0, e2))), inference(cnf_transformation, [], [f1])).
fof(f1153, plain, (spl24_106 | spl24_107 | spl24_108 | spl24_109 | spl24_110), inference(avatar_split_clause, [], [f299, f1150, f1146, f1142, f1138, f1134])).
fof(f299, plain, ((e4 = op(e0, e3)) | (e3 = op(e0, e3)) | (e2 = op(e0, e3)) | (e1 = op(e0, e3)) | (e0 = op(e0, e3))), inference(cnf_transformation, [], [f1])).
fof(f1132, plain, (spl24_101 | spl24_102 | spl24_103 | spl24_104 | spl24_105), inference(avatar_split_clause, [], [f300, f1129, f1125, f1121, f1117, f1113])).
fof(f300, plain, ((e4 = op(e0, e4)) | (e3 = op(e0, e4)) | (e2 = op(e0, e4)) | (e1 = op(e0, e4)) | (e0 = op(e0, e4))), inference(cnf_transformation, [], [f1])).
fof(f1111, plain, (spl24_96 | spl24_97 | spl24_98 | spl24_99 | spl24_100), inference(avatar_split_clause, [], [f301, f1108, f1104, f1100, f1096, f1092])).
fof(f301, plain, ((e4 = op(e1, e0)) | (e3 = op(e1, e0)) | (e2 = op(e1, e0)) | (e1 = op(e1, e0)) | (e0 = op(e1, e0))), inference(cnf_transformation, [], [f1])).
fof(f1090, plain, (spl24_91 | spl24_92 | spl24_93 | spl24_94 | spl24_95), inference(avatar_split_clause, [], [f302, f1087, f1083, f1079, f1075, f1071])).
fof(f302, plain, ((e4 = op(e1, e1)) | (e3 = op(e1, e1)) | (e2 = op(e1, e1)) | (e1 = op(e1, e1)) | (e0 = op(e1, e1))), inference(cnf_transformation, [], [f1])).
fof(f1069, plain, (spl24_86 | spl24_87 | spl24_88 | spl24_89 | spl24_90), inference(avatar_split_clause, [], [f303, f1066, f1062, f1058, f1054, f1050])).
fof(f303, plain, ((e4 = op(e1, e2)) | (e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (e0 = op(e1, e2))), inference(cnf_transformation, [], [f1])).
fof(f1048, plain, (spl24_81 | spl24_82 | spl24_83 | spl24_84 | spl24_85), inference(avatar_split_clause, [], [f304, f1045, f1041, f1037, f1033, f1029])).
fof(f304, plain, ((e4 = op(e1, e3)) | (e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (e0 = op(e1, e3))), inference(cnf_transformation, [], [f1])).
fof(f1027, plain, (spl24_76 | spl24_77 | spl24_78 | spl24_79 | spl24_80), inference(avatar_split_clause, [], [f305, f1024, f1020, f1016, f1012, f1008])).
fof(f305, plain, ((e4 = op(e1, e4)) | (e3 = op(e1, e4)) | (e2 = op(e1, e4)) | (e1 = op(e1, e4)) | (e0 = op(e1, e4))), inference(cnf_transformation, [], [f1])).
fof(f1006, plain, (spl24_71 | spl24_72 | spl24_73 | spl24_74 | spl24_75), inference(avatar_split_clause, [], [f306, f1003, f999, f995, f991, f987])).
fof(f306, plain, ((e4 = op(e2, e0)) | (e3 = op(e2, e0)) | (e2 = op(e2, e0)) | (e1 = op(e2, e0)) | (e0 = op(e2, e0))), inference(cnf_transformation, [], [f1])).
fof(f985, plain, (spl24_66 | spl24_67 | spl24_68 | spl24_69 | spl24_70), inference(avatar_split_clause, [], [f307, f982, f978, f974, f970, f966])).
fof(f307, plain, ((e4 = op(e2, e1)) | (e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))), inference(cnf_transformation, [], [f1])).
fof(f964, plain, (spl24_61 | spl24_62 | spl24_63 | spl24_64 | spl24_65), inference(avatar_split_clause, [], [f308, f961, f957, f953, f949, f945])).
fof(f308, plain, ((e4 = op(e2, e2)) | (e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (e0 = op(e2, e2))), inference(cnf_transformation, [], [f1])).
fof(f943, plain, (spl24_56 | spl24_57 | spl24_58 | spl24_59 | spl24_60), inference(avatar_split_clause, [], [f309, f940, f936, f932, f928, f924])).
fof(f309, plain, ((e4 = op(e2, e3)) | (e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))), inference(cnf_transformation, [], [f1])).
fof(f922, plain, (spl24_51 | spl24_52 | spl24_53 | spl24_54 | spl24_55), inference(avatar_split_clause, [], [f310, f919, f915, f911, f907, f903])).
fof(f310, plain, ((e4 = op(e2, e4)) | (e3 = op(e2, e4)) | (e2 = op(e2, e4)) | (e1 = op(e2, e4)) | (e0 = op(e2, e4))), inference(cnf_transformation, [], [f1])).
fof(f901, plain, (spl24_46 | spl24_47 | spl24_48 | spl24_49 | spl24_50), inference(avatar_split_clause, [], [f311, f898, f894, f890, f886, f882])).
fof(f311, plain, ((e4 = op(e3, e0)) | (e3 = op(e3, e0)) | (e2 = op(e3, e0)) | (e1 = op(e3, e0)) | (e0 = op(e3, e0))), inference(cnf_transformation, [], [f1])).
fof(f880, plain, (spl24_41 | spl24_42 | spl24_43 | spl24_44 | spl24_45), inference(avatar_split_clause, [], [f312, f877, f873, f869, f865, f861])).
fof(f312, plain, ((e4 = op(e3, e1)) | (e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (e0 = op(e3, e1))), inference(cnf_transformation, [], [f1])).
fof(f859, plain, (spl24_36 | spl24_37 | spl24_38 | spl24_39 | spl24_40), inference(avatar_split_clause, [], [f313, f856, f852, f848, f844, f840])).
fof(f313, plain, ((e4 = op(e3, e2)) | (e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (e0 = op(e3, e2))), inference(cnf_transformation, [], [f1])).
fof(f838, plain, (spl24_31 | spl24_32 | spl24_33 | spl24_34 | spl24_35), inference(avatar_split_clause, [], [f314, f835, f831, f827, f823, f819])).
fof(f314, plain, ((e4 = op(e3, e3)) | (e3 = op(e3, e3)) | (e2 = op(e3, e3)) | (e1 = op(e3, e3)) | (e0 = op(e3, e3))), inference(cnf_transformation, [], [f1])).
fof(f817, plain, (spl24_26 | spl24_27 | spl24_28 | spl24_29 | spl24_30), inference(avatar_split_clause, [], [f315, f814, f810, f806, f802, f798])).
fof(f315, plain, ((e4 = op(e3, e4)) | (e3 = op(e3, e4)) | (e2 = op(e3, e4)) | (e1 = op(e3, e4)) | (e0 = op(e3, e4))), inference(cnf_transformation, [], [f1])).
fof(f796, plain, (spl24_21 | spl24_22 | spl24_23 | spl24_24 | spl24_25), inference(avatar_split_clause, [], [f316, f793, f789, f785, f781, f777])).
fof(f316, plain, ((e4 = op(e4, e0)) | (e3 = op(e4, e0)) | (e2 = op(e4, e0)) | (e1 = op(e4, e0)) | (e0 = op(e4, e0))), inference(cnf_transformation, [], [f1])).
fof(f775, plain, (spl24_16 | spl24_17 | spl24_18 | spl24_19 | spl24_20), inference(avatar_split_clause, [], [f317, f772, f768, f764, f760, f756])).
fof(f317, plain, ((e4 = op(e4, e1)) | (e3 = op(e4, e1)) | (e2 = op(e4, e1)) | (e1 = op(e4, e1)) | (e0 = op(e4, e1))), inference(cnf_transformation, [], [f1])).
fof(f754, plain, (spl24_11 | spl24_12 | spl24_13 | spl24_14 | spl24_15), inference(avatar_split_clause, [], [f318, f751, f747, f743, f739, f735])).
fof(f318, plain, ((e4 = op(e4, e2)) | (e3 = op(e4, e2)) | (e2 = op(e4, e2)) | (e1 = op(e4, e2)) | (e0 = op(e4, e2))), inference(cnf_transformation, [], [f1])).
fof(f733, plain, (spl24_6 | spl24_7 | spl24_8 | spl24_9 | spl24_10), inference(avatar_split_clause, [], [f319, f730, f726, f722, f718, f714])).
fof(f319, plain, ((e4 = op(e4, e3)) | (e3 = op(e4, e3)) | (e2 = op(e4, e3)) | (e1 = op(e4, e3)) | (e0 = op(e4, e3))), inference(cnf_transformation, [], [f1])).
fof(f712, plain, (spl24_1 | spl24_2 | spl24_3 | spl24_4 | spl24_5), inference(avatar_split_clause, [], [f320, f709, f705, f701, f697, f693])).
fof(f320, plain, ((e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | (e1 = op(e4, e4)) | (e0 = op(e4, e4))), inference(cnf_transformation, [], [f1])).