fof(f4463, plain, $false, inference(avatar_sat_refutation, [], [f705, f722, f739, f756, f773, f790, f807, f824, f841, f858, f875, f892, f909, f926, f943, f960, f961, f962, f963, f964, f965, f966, f967, f968, f969, f970, f971, f972, f973, f975, f976, f977, f978, f979, f980, f981, f982, f983, f984, f985, f986, f987, f988, f989, f990, f991, f992, f997, f998, f999, f1000, f1001, f1006, f1007, f1008, f1009, f1010, f1015, f1016, f1017, f1018, f1019, f1024, f1025, f1026, f1027, f1028, f1033, f1034, f1035, f1037, f1042, f1043, f1044, f1045, f1046, f1051, f1052, f1053, f1054, f1055, f1060, f1061, f1062, f1063, f1064, f1069, f1070, f1071, f1072, f1073, f1078, f1079, f1080, f1082, f1087, f1088, f1089, f1090, f1091, f1096, f1097, f1098, f1099, f1100, f1105, f1106, f1107, f1108, f1109, f1114, f1115, f1116, f1117, f1118, f1123, f1124, f1125, f1127, f1134, f1135, f1142, f1143, f1149, f1150, f1151, f1157, f1158, f1159, f1164, f1165, f1167, f1172, f1173, f1174, f1175, f1180, f1181, f1182, f1183, f1189, f1190, f1196, f1197, f1198, f1199, f1204, f1205, f1212, f1213, f1214, f1215, f1221, f1222, f1223, f1228, f1229, f1230, f1231, f1236, f1237, f1238, f1239, f1244, f1245, f1252, f1253, f1261, f1262, f1268, f1269, f1270, f1271, f1276, f1277, f1278, f1279, f1286, f1287, f1294, f1295, f1301, f1302, f1303, f1310, f1311, f1316, f1317, f1318, f1319, f1325, f1326, f1332, f1333, f1340, f1341, f1342, f1343, f1348, f1349, f1350, f1351, f1357, f1358, f1359, f1364, f1365, f1366, f1367, f1372, f1373, f1375, f1380, f1381, f1388, f1389, f1390, f1391, f1397, f1398, f1404, f1405, f1406, f1407, f1412, f1413, f1414, f1415, f1420, f1421, f1429, f1430, f1436, f1437, f1438, f1439, f1446, f1447, f1454, f1455, f1462, f1463, f1470, f1471, f1476, f1477, f1478, f1479, f1484, f1485, f1486, f1487, f1493, f1494, f1495, f1500, f1501, f1503, f1508, f1509, f1516, f1517, f1518, f1519, f1524, f1525, f1526, f1527, f1533, f1534, f1535, f1540, f1541, f1542, f1543, f1548, f1549, f1556, f1557, f1558, f1559, f1565, f1566, f1572, f1573, f1574, f1575, f1580, f1581, f1582, f1583, f1588, f1589, f1597, f1598, f1599, f1606, f1607, f1614, f1615, f1622, f1623, f1630, f1631, f1632, f1633, f1634, f1636, f1639, f1640, f1649, f1654, f1663, f1668, f1677, f1682, f1691, f1696, f1701, f1706, f1711, f1716, f1721, f1726, f1731, f1736, f1741, f1742, f1747, f1748, f1753, f1754, f1759, f1760, f1769, f1770, f1772, f1773, f1789, f1792, f1802, f1809, f1810, f1813, f1822, f1824, f1825, f1829, f1835, f1841, f1844, f1850, f1855, f1867, f1868, f1872, f1875, f1887, f1900, f1902, f1914, f1917, f1949, f1953, f1957, f1963, f1967, f1970, f1975, f1983, f1987, f1997, f2000, f2003, f2009, f2026, f2033, f2044, f2053, f2105, f2116, f2130, f2138, f2148, f2151, f2182, f2185, f2195, f2239, f2242, f2249, f2266, f2273, f2290, f2311, f2321, f2346, f2349, f2361, f2375, f2382, f2406, f2411, f2416, f2423, f2433, f2436, f2445, f2467, f2477, f2484, f2489, f2499, f2539, f2548, f2559, f2614, f2616, f2621, f2627, f2639, f2645, f2656, f2688, f2696, f2699, f2727, f2729, f2759, f2794, f2805, f2812, f2816, f2823, f2833, f2839, f2857, f2865, f2877, f2905, f2912, f2925, f2930, f2950, f2951, f2959, f2976, f2997, f2999, f3003, f3006, f3008, f3015, f3025, f3029, f3033, f3038, f3045, f3054, f3058, f3066, f3079, f3131, f3144, f3154, f3167, f3179, f3202, f3223, f3233, f3257, f3283, f3291, f3322, f3341, f3347, f3352, f3360, f3362, f3423, f3431, f3438, f3439, f3443, f3444, f3452, f3464, f3469, f3476, f3515, f3521, f3526, f3531, f3541, f3546, f3559, f3571, f3617, f3623, f3631, f3638, f3643, f3657, f3659, f3664, f3668, f3677, f3680, f3690, f3700, f3702, f3716, f3744, f3745, f3749, f3754, f3759, f3783, f3798, f3799, f3800, f3806, f3814, f3815, f3828, f3829, f3839, f3845, f3863, f3885, f3889, f3898, f3899, f3904, f3906, f3910, f3913, f3915, f3919, f3922, f3926, f3931, f3933, f3939, f3941, f3943, f3945, f3947, f3949, f3951, f3953, f3961, f3993, f3994, f4025, f4039, f4048, f4064, f4065, f4080, f4081, f4082, f4097, f4100, f4115, f4116, f4124, f4132, f4136, f4145, f4154, f4171, f4187, f4188, f4189, f4197, f4203, f4211, f4219, f4226, f4233, f4238, f4244, f4258, f4261, f4263, f4265, f4280, f4283, f4299, f4300, f4316, f4317, f4323, f4332, f4333, f4342, f4343, f4348, f4354, f4359, f4370, f4375, f4376, f4384, f4411, f4412, f4421, f4422, f4430, f4437, f4447, f4450, f4457, f4462])).
fof(f4462, plain, (~ spl78_2 | ~ spl78_4), inference(avatar_contradiction_clause, [], [f4461])).
fof(f4461, plain, ($false | (~ spl78_2 | ~ spl78_4)), inference(subsumption_resolution, [], [f4460, f311])).
fof(f311, plain, ~ (e1 = e3), inference(cnf_transformation, [], [f4])).
fof(f4, plain, (~ (e2 = e3) & ~ (e1 = e3) & ~ (e1 = e2) & ~ (e0 = e3) & ~ (e0 = e2) & ~ (e0 = e1)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG160+1.p', ax4)).
fof(f4460, plain, ((e1 = e3) | (~ spl78_2 | ~ spl78_4)), inference(backward_demodulation, [], [f704, f696])).
fof(f696, plain, ((e1 = op(e3, e3)) | ~ spl78_2), inference(avatar_component_clause, [], [f694])).
fof(f694, plain, (spl78_2 <=> (e1 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl78_2])])).
fof(f704, plain, ((e3 = op(e3, e3)) | ~ spl78_4), inference(avatar_component_clause, [], [f702])).
fof(f702, plain, (spl78_4 <=> (e3 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl78_4])])).
fof(f4457, plain, (~ spl78_5 | ~ spl78_9), inference(avatar_split_clause, [], [f4453, f724, f707])).
fof(f707, plain, (spl78_5 <=> (e0 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl78_5])])).
fof(f724, plain, (spl78_9 <=> (e0 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl78_9])])).
fof(f4453, plain, (~ (e0 = op(e3, e2)) | ~ spl78_9), inference(backward_demodulation, [], [f304, f726])).
fof(f726, plain, ((e0 = op(e3, e1)) | ~ spl78_9), inference(avatar_component_clause, [], [f724])).
fof(f304, plain, ~ (op(e3, e1) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f3, plain, (~ (op(e3, e2) = op(e3, e3)) & ~ (op(e3, e1) = op(e3, e3)) & ~ (op(e3, e1) = op(e3, e2)) & ~ (op(e3, e0) = op(e3, e3)) & ~ (op(e3, e0) = op(e3, e2)) & ~ (op(e3, e0) = op(e3, e1)) & ~ (op(e2, e2) = op(e2, e3)) & ~ (op(e2, e1) = op(e2, e3)) & ~ (op(e2, e1) = op(e2, e2)) & ~ (op(e2, e0) = op(e2, e3)) & ~ (op(e2, e0) = op(e2, e2)) & ~ (op(e2, e0) = op(e2, e1)) & ~ (op(e1, e2) = op(e1, e3)) & ~ (op(e1, e1) = op(e1, e3)) & ~ (op(e1, e1) = op(e1, e2)) & ~ (op(e1, e0) = op(e1, e3)) & ~ (op(e1, e0) = op(e1, e2)) & ~ (op(e1, e0) = op(e1, e1)) & ~ (op(e0, e2) = op(e0, e3)) & ~ (op(e0, e1) = op(e0, e3)) & ~ (op(e0, e1) = op(e0, e2)) & ~ (op(e0, e0) = op(e0, e3)) & ~ (op(e0, e0) = op(e0, e2)) & ~ (op(e0, e0) = op(e0, e1)) & ~ (op(e2, e3) = op(e3, e3)) & ~ (op(e1, e3) = op(e3, e3)) & ~ (op(e1, e3) = op(e2, e3)) & ~ (op(e0, e3) = op(e3, e3)) & ~ (op(e0, e3) = op(e2, e3)) & ~ (op(e0, e3) = op(e1, e3)) & ~ (op(e2, e2) = op(e3, e2)) & ~ (op(e1, e2) = op(e3, e2)) & ~ (op(e1, e2) = op(e2, e2)) & ~ (op(e0, e2) = op(e3, e2)) & ~ (op(e0, e2) = op(e2, e2)) & ~ (op(e0, e2) = op(e1, e2)) & ~ (op(e2, e1) = op(e3, e1)) & ~ (op(e1, e1) = op(e3, e1)) & ~ (op(e1, e1) = op(e2, e1)) & ~ (op(e0, e1) = op(e3, e1)) & ~ (op(e0, e1) = op(e2, e1)) & ~ (op(e0, e1) = op(e1, e1)) & ~ (op(e2, e0) = op(e3, e0)) & ~ (op(e1, e0) = op(e3, e0)) & ~ (op(e1, e0) = op(e2, e0)) & ~ (op(e0, e0) = op(e3, e0)) & ~ (op(e0, e0) = op(e2, e0)) & ~ (op(e0, e0) = op(e1, e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG160+1.p', ax3)).
fof(f4450, plain, (~ spl78_11 | ~ spl78_27), inference(avatar_split_clause, [], [f4449, f800, f732])).
fof(f732, plain, (spl78_11 <=> (e2 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl78_11])])).
fof(f800, plain, (spl78_27 <=> (e2 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl78_27])])).
fof(f4449, plain, (~ (e2 = op(e3, e1)) | ~ spl78_27), inference(backward_demodulation, [], [f270, f802])).
fof(f802, plain, ((e2 = op(e2, e1)) | ~ spl78_27), inference(avatar_component_clause, [], [f800])).
fof(f270, plain, ~ (op(e2, e1) = op(e3, e1)), inference(cnf_transformation, [], [f3])).
fof(f4447, plain, (~ spl78_26 | ~ spl78_30), inference(avatar_split_clause, [], [f4443, f813, f796])).
fof(f796, plain, (spl78_26 <=> (e1 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl78_26])])).
fof(f813, plain, (spl78_30 <=> (e1 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl78_30])])).
fof(f4443, plain, (~ (e1 = op(e2, e1)) | ~ spl78_30), inference(backward_demodulation, [], [f295, f815])).
fof(f815, plain, ((e1 = op(e2, e0)) | ~ spl78_30), inference(avatar_component_clause, [], [f813])).
fof(f295, plain, ~ (op(e2, e0) = op(e2, e1)), inference(cnf_transformation, [], [f3])).
fof(f4437, plain, (~ spl78_41 | ~ spl78_45), inference(avatar_split_clause, [], [f4433, f877, f860])).
fof(f860, plain, (spl78_41 <=> (e0 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl78_41])])).
fof(f877, plain, (spl78_45 <=> (e0 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl78_45])])).
fof(f4433, plain, (~ (e0 = op(e1, e1)) | ~ spl78_45), inference(backward_demodulation, [], [f289, f879])).
fof(f879, plain, ((e0 = op(e1, e0)) | ~ spl78_45), inference(avatar_component_clause, [], [f877])).
fof(f289, plain, ~ (op(e1, e0) = op(e1, e1)), inference(cnf_transformation, [], [f3])).
fof(f4430, plain, (~ spl78_5 | ~ spl78_53), inference(avatar_split_clause, [], [f4423, f911, f707])).
fof(f911, plain, (spl78_53 <=> (e0 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl78_53])])).
fof(f4423, plain, (~ (e0 = op(e3, e2)) | ~ spl78_53), inference(backward_demodulation, [], [f273, f913])).
fof(f913, plain, ((e0 = op(e0, e2)) | ~ spl78_53), inference(avatar_component_clause, [], [f911])).
fof(f273, plain, ~ (op(e0, e2) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f4422, plain, (~ spl78_50 | ~ spl78_58), inference(avatar_split_clause, [], [f4418, f932, f898])).
fof(f898, plain, (spl78_50 <=> (e1 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl78_50])])).
fof(f932, plain, (spl78_58 <=> (e1 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl78_58])])).
fof(f4418, plain, (~ (e1 = op(e0, e3)) | ~ spl78_58), inference(backward_demodulation, [], [f287, f934])).
fof(f934, plain, ((e1 = op(e0, e1)) | ~ spl78_58), inference(avatar_component_clause, [], [f932])).
fof(f287, plain, ~ (op(e0, e1) = op(e0, e3)), inference(cnf_transformation, [], [f3])).
fof(f4421, plain, (~ spl78_26 | ~ spl78_58), inference(avatar_split_clause, [], [f4415, f932, f796])).
fof(f4415, plain, (~ (e1 = op(e2, e1)) | ~ spl78_58), inference(backward_demodulation, [], [f266, f934])).
fof(f266, plain, ~ (op(e0, e1) = op(e2, e1)), inference(cnf_transformation, [], [f3])).
fof(f4412, plain, (~ spl78_55 | ~ spl78_63), inference(avatar_split_clause, [], [f4403, f953, f919])).
fof(f919, plain, (spl78_55 <=> (e2 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl78_55])])).
fof(f953, plain, (spl78_63 <=> (op(e0, e0) = e2)), introduced(avatar_definition, [new_symbols(naming, [spl78_63])])).
fof(f4403, plain, (~ (e2 = op(e0, e2)) | ~ spl78_63), inference(backward_demodulation, [], [f284, f955])).
fof(f955, plain, ((op(e0, e0) = e2) | ~ spl78_63), inference(avatar_component_clause, [], [f953])).
fof(f284, plain, ~ (op(e0, e0) = op(e0, e2)), inference(cnf_transformation, [], [f3])).
fof(f4411, plain, (~ spl78_31 | ~ spl78_63), inference(avatar_split_clause, [], [f4400, f953, f817])).
fof(f817, plain, (spl78_31 <=> (e2 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl78_31])])).
fof(f4400, plain, (~ (e2 = op(e2, e0)) | ~ spl78_63), inference(backward_demodulation, [], [f260, f955])).
fof(f260, plain, ~ (op(e0, e0) = op(e2, e0)), inference(cnf_transformation, [], [f3])).
fof(f4384, plain, (~ spl78_52 | ~ spl78_4), inference(avatar_split_clause, [], [f4058, f702, f906])).
fof(f906, plain, (spl78_52 <=> (e3 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl78_52])])).
fof(f4058, plain, (~ (e3 = op(e0, e3)) | ~ spl78_4), inference(forward_demodulation, [], [f279, f704])).
fof(f279, plain, ~ (op(e0, e3) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f4376, plain, (~ spl78_16 | ~ spl78_4), inference(avatar_split_clause, [], [f4050, f702, f753])).
fof(f753, plain, (spl78_16 <=> (e3 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl78_16])])).
fof(f4050, plain, (~ (e3 = op(e3, e0)) | ~ spl78_4), inference(forward_demodulation, [], [f303, f704])).
fof(f303, plain, ~ (op(e3, e0) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f4375, plain, (~ spl78_9 | spl78_162 | ~ spl78_165), inference(avatar_split_clause, [], [f4274, f1750, f1733, f724])).
fof(f1733, plain, (spl78_162 <=> (e0 = op(op(e1, op(e1, e1)), e1))), introduced(avatar_definition, [new_symbols(naming, [spl78_162])])).
fof(f1750, plain, (spl78_165 <=> (e3 = op(e1, op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl78_165])])).
fof(f4274, plain, (~ (e0 = op(e3, e1)) | (spl78_162 | ~ spl78_165)), inference(backward_demodulation, [], [f1735, f1751])).
fof(f1751, plain, ((e3 = op(e1, op(e1, e1))) | ~ spl78_165), inference(avatar_component_clause, [], [f1750])).
fof(f1735, plain, (~ (e0 = op(op(e1, op(e1, e1)), e1)) | spl78_162), inference(avatar_component_clause, [], [f1733])).
fof(f4370, plain, (~ spl78_6 | ~ spl78_14), inference(avatar_split_clause, [], [f4369, f745, f711])).
fof(f711, plain, (spl78_6 <=> (e1 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl78_6])])).
fof(f745, plain, (spl78_14 <=> (e1 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl78_14])])).
fof(f4369, plain, (~ (e1 = op(e3, e2)) | ~ spl78_14), inference(backward_demodulation, [], [f302, f747])).
fof(f747, plain, ((e1 = op(e3, e0)) | ~ spl78_14), inference(avatar_component_clause, [], [f745])).
fof(f302, plain, ~ (op(e3, e0) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f4359, plain, (~ spl78_18 | ~ spl78_26), inference(avatar_split_clause, [], [f4356, f796, f762])).
fof(f762, plain, (spl78_18 <=> (e1 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl78_18])])).
fof(f4356, plain, (~ (e1 = op(e2, e3)) | ~ spl78_26), inference(backward_demodulation, [], [f299, f798])).
fof(f798, plain, ((e1 = op(e2, e1)) | ~ spl78_26), inference(avatar_component_clause, [], [f796])).
fof(f299, plain, ~ (op(e2, e1) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f4354, plain, (~ spl78_23 | ~ spl78_31), inference(avatar_split_clause, [], [f4351, f817, f783])).
fof(f783, plain, (spl78_23 <=> (e2 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl78_23])])).
fof(f4351, plain, (~ (e2 = op(e2, e2)) | ~ spl78_31), inference(backward_demodulation, [], [f296, f819])).
fof(f819, plain, ((e2 = op(e2, e0)) | ~ spl78_31), inference(avatar_component_clause, [], [f817])).
fof(f296, plain, ~ (op(e2, e0) = op(e2, e2)), inference(cnf_transformation, [], [f3])).
fof(f4348, plain, (~ spl78_6 | ~ spl78_38), inference(avatar_split_clause, [], [f4345, f847, f711])).
fof(f847, plain, (spl78_38 <=> (e1 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl78_38])])).
fof(f4345, plain, (~ (e1 = op(e3, e2)) | ~ spl78_38), inference(backward_demodulation, [], [f275, f849])).
fof(f849, plain, ((e1 = op(e1, e2)) | ~ spl78_38), inference(avatar_component_clause, [], [f847])).
fof(f275, plain, ~ (op(e1, e2) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f4343, plain, (~ spl78_37 | ~ spl78_41), inference(avatar_split_clause, [], [f4335, f860, f843])).
fof(f843, plain, (spl78_37 <=> (e0 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl78_37])])).
fof(f4335, plain, (~ (e0 = op(e1, e2)) | ~ spl78_41), inference(backward_demodulation, [], [f292, f862])).
fof(f862, plain, ((e0 = op(e1, e1)) | ~ spl78_41), inference(avatar_component_clause, [], [f860])).
fof(f292, plain, ~ (op(e1, e1) = op(e1, e2)), inference(cnf_transformation, [], [f3])).
fof(f4342, plain, (~ spl78_25 | ~ spl78_41), inference(avatar_split_clause, [], [f4334, f860, f792])).
fof(f792, plain, (spl78_25 <=> (e0 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl78_25])])).
fof(f4334, plain, (~ (e0 = op(e2, e1)) | ~ spl78_41), inference(backward_demodulation, [], [f268, f862])).
fof(f268, plain, ~ (op(e1, e1) = op(e2, e1)), inference(cnf_transformation, [], [f3])).
fof(f4333, plain, (~ spl78_44 | ~ spl78_48), inference(avatar_split_clause, [], [f4330, f889, f872])).
fof(f872, plain, (spl78_44 <=> (e3 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl78_44])])).
fof(f889, plain, (spl78_48 <=> (e3 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl78_48])])).
fof(f4330, plain, (~ (e3 = op(e1, e1)) | ~ spl78_48), inference(backward_demodulation, [], [f289, f891])).
fof(f891, plain, ((e3 = op(e1, e0)) | ~ spl78_48), inference(avatar_component_clause, [], [f889])).
fof(f4332, plain, (~ spl78_32 | ~ spl78_48), inference(avatar_split_clause, [], [f4328, f889, f821])).
fof(f821, plain, (spl78_32 <=> (e3 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl78_32])])).
fof(f4328, plain, (~ (e3 = op(e2, e0)) | ~ spl78_48), inference(backward_demodulation, [], [f262, f891])).
fof(f262, plain, ~ (op(e1, e0) = op(e2, e0)), inference(cnf_transformation, [], [f3])).
fof(f4323, plain, (~ spl78_23 | ~ spl78_55), inference(avatar_split_clause, [], [f4319, f919, f783])).
fof(f4319, plain, (~ (e2 = op(e2, e2)) | ~ spl78_55), inference(backward_demodulation, [], [f272, f921])).
fof(f921, plain, ((e2 = op(e0, e2)) | ~ spl78_55), inference(avatar_component_clause, [], [f919])).
fof(f272, plain, ~ (op(e0, e2) = op(e2, e2)), inference(cnf_transformation, [], [f3])).
fof(f4317, plain, (~ spl78_56 | ~ spl78_60), inference(avatar_split_clause, [], [f4314, f940, f923])).
fof(f923, plain, (spl78_56 <=> (e3 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl78_56])])).
fof(f940, plain, (spl78_60 <=> (e3 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl78_60])])).
fof(f4314, plain, (~ (e3 = op(e0, e2)) | ~ spl78_60), inference(backward_demodulation, [], [f286, f942])).
fof(f942, plain, ((e3 = op(e0, e1)) | ~ spl78_60), inference(avatar_component_clause, [], [f940])).
fof(f286, plain, ~ (op(e0, e1) = op(e0, e2)), inference(cnf_transformation, [], [f3])).
fof(f4316, plain, (~ spl78_44 | ~ spl78_60), inference(avatar_split_clause, [], [f4312, f940, f872])).
fof(f4312, plain, (~ (e3 = op(e1, e1)) | ~ spl78_60), inference(backward_demodulation, [], [f265, f942])).
fof(f265, plain, ~ (op(e0, e1) = op(e1, e1)), inference(cnf_transformation, [], [f3])).
fof(f4300, plain, (~ spl78_49 | ~ spl78_61), inference(avatar_split_clause, [], [f4289, f945, f894])).
fof(f894, plain, (spl78_49 <=> (e0 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl78_49])])).
fof(f945, plain, (spl78_61 <=> (e0 = op(e0, e0))), introduced(avatar_definition, [new_symbols(naming, [spl78_61])])).
fof(f4289, plain, (~ (e0 = op(e0, e3)) | ~ spl78_61), inference(backward_demodulation, [], [f285, f947])).
fof(f947, plain, ((e0 = op(e0, e0)) | ~ spl78_61), inference(avatar_component_clause, [], [f945])).
fof(f285, plain, ~ (op(e0, e0) = op(e0, e3)), inference(cnf_transformation, [], [f3])).
fof(f4299, plain, (~ spl78_13 | ~ spl78_61), inference(avatar_split_clause, [], [f4286, f945, f741])).
fof(f741, plain, (spl78_13 <=> (e0 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl78_13])])).
fof(f4286, plain, (~ (e0 = op(e3, e0)) | ~ spl78_61), inference(backward_demodulation, [], [f261, f947])).
fof(f261, plain, ~ (op(e0, e0) = op(e3, e0)), inference(cnf_transformation, [], [f3])).
fof(f4283, plain, (~ spl78_11 | spl78_145 | ~ spl78_165), inference(avatar_contradiction_clause, [], [f4282])).
fof(f4282, plain, ($false | (~ spl78_11 | spl78_145 | ~ spl78_165)), inference(subsumption_resolution, [], [f4281, f734])).
fof(f734, plain, ((e2 = op(e3, e1)) | ~ spl78_11), inference(avatar_component_clause, [], [f732])).
fof(f4281, plain, (~ (e2 = op(e3, e1)) | (spl78_145 | ~ spl78_165)), inference(forward_demodulation, [], [f1653, f1751])).
fof(f1653, plain, (~ (e2 = op(op(e1, op(e1, e1)), e1)) | spl78_145), inference(avatar_component_clause, [], [f1651])).
fof(f1651, plain, (spl78_145 <=> (e2 = op(op(e1, op(e1, e1)), e1))), introduced(avatar_definition, [new_symbols(naming, [spl78_145])])).
fof(f4280, plain, (~ spl78_56 | ~ spl78_146 | spl78_147), inference(avatar_split_clause, [], [f4277, f1660, f1656, f923])).
fof(f1656, plain, (spl78_146 <=> (e0 = op(e2, op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl78_146])])).
fof(f1660, plain, (spl78_147 <=> (e3 = op(op(e2, op(e2, e2)), e2))), introduced(avatar_definition, [new_symbols(naming, [spl78_147])])).
fof(f4277, plain, (~ (e3 = op(e0, e2)) | (~ spl78_146 | spl78_147)), inference(backward_demodulation, [], [f1662, f1657])).
fof(f1657, plain, ((e0 = op(e2, op(e2, e2))) | ~ spl78_146), inference(avatar_component_clause, [], [f1656])).
fof(f1662, plain, (~ (e3 = op(op(e2, op(e2, e2)), e2)) | spl78_147), inference(avatar_component_clause, [], [f1660])).
fof(f4265, plain, (~ spl78_19 | ~ spl78_35), inference(avatar_split_clause, [], [f4264, f834, f766])).
fof(f766, plain, (spl78_19 <=> (e2 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl78_19])])).
fof(f834, plain, (spl78_35 <=> (e2 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl78_35])])).
fof(f4264, plain, (~ (e2 = op(e2, e3)) | ~ spl78_35), inference(forward_demodulation, [], [f280, f836])).
fof(f836, plain, ((e2 = op(e1, e3)) | ~ spl78_35), inference(avatar_component_clause, [], [f834])).
fof(f280, plain, ~ (op(e1, e3) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f4263, plain, (~ spl78_15 | ~ spl78_11), inference(avatar_split_clause, [], [f4262, f732, f749])).
fof(f749, plain, (spl78_15 <=> (e2 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl78_15])])).
fof(f4262, plain, (~ (e2 = op(e3, e0)) | ~ spl78_11), inference(forward_demodulation, [], [f301, f734])).
fof(f301, plain, ~ (op(e3, e0) = op(e3, e1)), inference(cnf_transformation, [], [f3])).
fof(f4261, plain, (~ spl78_7 | ~ spl78_11), inference(avatar_split_clause, [], [f4260, f732, f715])).
fof(f715, plain, (spl78_7 <=> (e2 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl78_7])])).
fof(f4260, plain, (~ (e2 = op(e3, e2)) | ~ spl78_11), inference(forward_demodulation, [], [f304, f734])).
fof(f4258, plain, (~ spl78_9 | ~ spl78_13), inference(avatar_split_clause, [], [f4256, f741, f724])).
fof(f4256, plain, (~ (e0 = op(e3, e1)) | ~ spl78_13), inference(backward_demodulation, [], [f301, f743])).
fof(f743, plain, ((e0 = op(e3, e0)) | ~ spl78_13), inference(avatar_component_clause, [], [f741])).
fof(f4244, plain, (~ spl78_9 | ~ spl78_25), inference(avatar_split_clause, [], [f4239, f792, f724])).
fof(f4239, plain, (~ (e0 = op(e3, e1)) | ~ spl78_25), inference(backward_demodulation, [], [f270, f794])).
fof(f794, plain, ((e0 = op(e2, e1)) | ~ spl78_25), inference(avatar_component_clause, [], [f792])).
fof(f4238, plain, (~ spl78_24 | ~ spl78_32), inference(avatar_split_clause, [], [f4237, f821, f787])).
fof(f787, plain, (spl78_24 <=> (e3 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl78_24])])).
fof(f4237, plain, (~ (e3 = op(e2, e2)) | ~ spl78_32), inference(backward_demodulation, [], [f296, f823])).
fof(f823, plain, ((e3 = op(e2, e0)) | ~ spl78_32), inference(avatar_component_clause, [], [f821])).
fof(f4233, plain, (~ spl78_33 | ~ spl78_37), inference(avatar_split_clause, [], [f4229, f843, f826])).
fof(f826, plain, (spl78_33 <=> (e0 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl78_33])])).
fof(f4229, plain, (~ (e0 = op(e1, e3)) | ~ spl78_37), inference(backward_demodulation, [], [f294, f845])).
fof(f845, plain, ((e0 = op(e1, e2)) | ~ spl78_37), inference(avatar_component_clause, [], [f843])).
fof(f294, plain, ~ (op(e1, e2) = op(e1, e3)), inference(cnf_transformation, [], [f3])).
fof(f4226, plain, (~ spl78_33 | ~ spl78_44 | spl78_143), inference(avatar_split_clause, [], [f4224, f1642, f872, f826])).
fof(f1642, plain, (spl78_143 <=> (e0 = op(e1, op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl78_143])])).
fof(f4224, plain, (~ (e0 = op(e1, e3)) | (~ spl78_44 | spl78_143)), inference(backward_demodulation, [], [f1644, f874])).
fof(f874, plain, ((e3 = op(e1, e1)) | ~ spl78_44), inference(avatar_component_clause, [], [f872])).
fof(f1644, plain, (~ (e0 = op(e1, op(e1, e1))) | spl78_143), inference(avatar_component_clause, [], [f1642])).
fof(f4219, plain, (~ spl78_42 | ~ spl78_46), inference(avatar_split_clause, [], [f4214, f881, f864])).
fof(f864, plain, (spl78_42 <=> (e1 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl78_42])])).
fof(f881, plain, (spl78_46 <=> (e1 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl78_46])])).
fof(f4214, plain, (~ (e1 = op(e1, e1)) | ~ spl78_46), inference(backward_demodulation, [], [f289, f883])).
fof(f883, plain, ((e1 = op(e1, e0)) | ~ spl78_46), inference(avatar_component_clause, [], [f881])).
fof(f4211, plain, (~ spl78_33 | ~ spl78_49), inference(avatar_split_clause, [], [f4209, f894, f826])).
fof(f4209, plain, (~ (e0 = op(e1, e3)) | ~ spl78_49), inference(backward_demodulation, [], [f277, f896])).
fof(f896, plain, ((e0 = op(e0, e3)) | ~ spl78_49), inference(avatar_component_clause, [], [f894])).
fof(f277, plain, ~ (op(e0, e3) = op(e1, e3)), inference(cnf_transformation, [], [f3])).
fof(f4203, plain, (~ spl78_24 | ~ spl78_56), inference(avatar_split_clause, [], [f4199, f923, f787])).
fof(f4199, plain, (~ (e3 = op(e2, e2)) | ~ spl78_56), inference(backward_demodulation, [], [f272, f925])).
fof(f925, plain, ((e3 = op(e0, e2)) | ~ spl78_56), inference(avatar_component_clause, [], [f923])).
fof(f4197, plain, (~ spl78_42 | ~ spl78_58), inference(avatar_split_clause, [], [f4190, f932, f864])).
fof(f4190, plain, (~ (e1 = op(e1, e1)) | ~ spl78_58), inference(backward_demodulation, [], [f265, f934])).
fof(f4189, plain, (~ spl78_56 | ~ spl78_63 | spl78_164), inference(avatar_split_clause, [], [f4186, f1744, f953, f923])).
fof(f1744, plain, (spl78_164 <=> (e3 = op(e0, op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl78_164])])).
fof(f4186, plain, (~ (e3 = op(e0, e2)) | (~ spl78_63 | spl78_164)), inference(backward_demodulation, [], [f1746, f955])).
fof(f1746, plain, (~ (e3 = op(e0, op(e0, e0))) | spl78_164), inference(avatar_component_clause, [], [f1744])).
fof(f4188, plain, (~ spl78_51 | ~ spl78_63), inference(avatar_split_clause, [], [f4181, f953, f902])).
fof(f902, plain, (spl78_51 <=> (e2 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl78_51])])).
fof(f4181, plain, (~ (e2 = op(e0, e3)) | ~ spl78_63), inference(backward_demodulation, [], [f285, f955])).
fof(f4187, plain, (~ spl78_15 | ~ spl78_63), inference(avatar_split_clause, [], [f4178, f953, f749])).
fof(f4178, plain, (~ (e2 = op(e3, e0)) | ~ spl78_63), inference(backward_demodulation, [], [f261, f955])).
fof(f4171, plain, (~ spl78_25 | ~ spl78_161 | spl78_162), inference(avatar_split_clause, [], [f4169, f1733, f1728, f792])).
fof(f1728, plain, (spl78_161 <=> (e2 = op(e1, op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl78_161])])).
fof(f4169, plain, (~ (e0 = op(e2, e1)) | (~ spl78_161 | spl78_162)), inference(backward_demodulation, [], [f1735, f1729])).
fof(f1729, plain, ((e2 = op(e1, op(e1, e1))) | ~ spl78_161), inference(avatar_component_clause, [], [f1728])).
fof(f4154, plain, (~ spl78_7 | ~ spl78_15), inference(avatar_split_clause, [], [f4152, f749, f715])).
fof(f4152, plain, (~ (e2 = op(e3, e2)) | ~ spl78_15), inference(backward_demodulation, [], [f302, f751])).
fof(f751, plain, ((e2 = op(e3, e0)) | ~ spl78_15), inference(avatar_component_clause, [], [f749])).
fof(f4145, plain, (~ spl78_18 | ~ spl78_24 | spl78_155), inference(avatar_split_clause, [], [f4142, f1698, f787, f762])).
fof(f1698, plain, (spl78_155 <=> (e1 = op(e2, op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl78_155])])).
fof(f4142, plain, (~ (e1 = op(e2, e3)) | (~ spl78_24 | spl78_155)), inference(backward_demodulation, [], [f1700, f789])).
fof(f789, plain, ((e3 = op(e2, e2)) | ~ spl78_24), inference(avatar_component_clause, [], [f787])).
fof(f1700, plain, (~ (e1 = op(e2, op(e2, e2))) | spl78_155), inference(avatar_component_clause, [], [f1698])).
fof(f4136, plain, (~ spl78_19 | ~ spl78_27), inference(avatar_split_clause, [], [f4134, f800, f766])).
fof(f4134, plain, (~ (e2 = op(e2, e3)) | ~ spl78_27), inference(backward_demodulation, [], [f299, f802])).
fof(f4132, plain, (~ spl78_21 | ~ spl78_29), inference(avatar_split_clause, [], [f4127, f809, f775])).
fof(f775, plain, (spl78_21 <=> (e0 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl78_21])])).
fof(f809, plain, (spl78_29 <=> (e0 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl78_29])])).
fof(f4127, plain, (~ (e0 = op(e2, e2)) | ~ spl78_29), inference(backward_demodulation, [], [f296, f811])).
fof(f811, plain, ((e0 = op(e2, e0)) | ~ spl78_29), inference(avatar_component_clause, [], [f809])).
fof(f4124, plain, (~ spl78_7 | ~ spl78_39), inference(avatar_split_clause, [], [f4122, f851, f715])).
fof(f851, plain, (spl78_39 <=> (e2 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl78_39])])).
fof(f4122, plain, (~ (e2 = op(e3, e2)) | ~ spl78_39), inference(backward_demodulation, [], [f275, f853])).
fof(f853, plain, ((e2 = op(e1, e2)) | ~ spl78_39), inference(avatar_component_clause, [], [f851])).
fof(f4116, plain, (~ spl78_38 | ~ spl78_42), inference(avatar_split_clause, [], [f4108, f864, f847])).
fof(f4108, plain, (~ (e1 = op(e1, e2)) | ~ spl78_42), inference(backward_demodulation, [], [f292, f866])).
fof(f866, plain, ((e1 = op(e1, e1)) | ~ spl78_42), inference(avatar_component_clause, [], [f864])).
fof(f4115, plain, (~ spl78_26 | ~ spl78_42), inference(avatar_split_clause, [], [f4107, f864, f796])).
fof(f4107, plain, (~ (e1 = op(e2, e1)) | ~ spl78_42), inference(backward_demodulation, [], [f268, f866])).
fof(f4100, plain, (~ spl78_19 | ~ spl78_51), inference(avatar_split_clause, [], [f4098, f902, f766])).
fof(f4098, plain, (~ (e2 = op(e2, e3)) | ~ spl78_51), inference(backward_demodulation, [], [f278, f904])).
fof(f904, plain, ((e2 = op(e0, e3)) | ~ spl78_51), inference(avatar_component_clause, [], [f902])).
fof(f278, plain, ~ (op(e0, e3) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f4097, plain, (~ spl78_21 | ~ spl78_53), inference(avatar_split_clause, [], [f4091, f911, f775])).
fof(f4091, plain, (~ (e0 = op(e2, e2)) | ~ spl78_53), inference(backward_demodulation, [], [f272, f913])).
fof(f4082, plain, (~ spl78_59 | ~ spl78_62 | spl78_159), inference(avatar_split_clause, [], [f4078, f1718, f949, f936])).
fof(f936, plain, (spl78_59 <=> (e2 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl78_59])])).
fof(f949, plain, (spl78_62 <=> (op(e0, e0) = e1)), introduced(avatar_definition, [new_symbols(naming, [spl78_62])])).
fof(f1718, plain, (spl78_159 <=> (e2 = op(e0, op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl78_159])])).
fof(f4078, plain, (~ (e2 = op(e0, e1)) | (~ spl78_62 | spl78_159)), inference(backward_demodulation, [], [f1720, f951])).
fof(f951, plain, ((op(e0, e0) = e1) | ~ spl78_62), inference(avatar_component_clause, [], [f949])).
fof(f1720, plain, (~ (e2 = op(e0, op(e0, e0))) | spl78_159), inference(avatar_component_clause, [], [f1718])).
fof(f4081, plain, (~ spl78_50 | ~ spl78_62), inference(avatar_split_clause, [], [f4075, f949, f898])).
fof(f4075, plain, (~ (e1 = op(e0, e3)) | ~ spl78_62), inference(backward_demodulation, [], [f285, f951])).
fof(f4080, plain, (~ spl78_14 | ~ spl78_62), inference(avatar_split_clause, [], [f4072, f949, f745])).
fof(f4072, plain, (~ (e1 = op(e3, e0)) | ~ spl78_62), inference(backward_demodulation, [], [f261, f951])).
fof(f4065, plain, (~ spl78_14 | spl78_160 | ~ spl78_164), inference(avatar_split_clause, [], [f4063, f1744, f1723, f745])).
fof(f1723, plain, (spl78_160 <=> (e1 = op(op(e0, op(e0, e0)), e0))), introduced(avatar_definition, [new_symbols(naming, [spl78_160])])).
fof(f4063, plain, (~ (e1 = op(e3, e0)) | (spl78_160 | ~ spl78_164)), inference(backward_demodulation, [], [f1725, f1745])).
fof(f1745, plain, ((e3 = op(e0, op(e0, e0))) | ~ spl78_164), inference(avatar_component_clause, [], [f1744])).
fof(f1725, plain, (~ (e1 = op(op(e0, op(e0, e0)), e0)) | spl78_160), inference(avatar_component_clause, [], [f1723])).
fof(f4064, plain, (~ spl78_15 | spl78_154 | ~ spl78_164), inference(avatar_split_clause, [], [f4062, f1744, f1693, f749])).
fof(f1693, plain, (spl78_154 <=> (e2 = op(op(e0, op(e0, e0)), e0))), introduced(avatar_definition, [new_symbols(naming, [spl78_154])])).
fof(f4062, plain, (~ (e2 = op(e3, e0)) | (spl78_154 | ~ spl78_164)), inference(backward_demodulation, [], [f1695, f1745])).
fof(f1695, plain, (~ (e2 = op(op(e0, op(e0, e0)), e0)) | spl78_154), inference(avatar_component_clause, [], [f1693])).
fof(f4048, plain, (~ spl78_3 | ~ spl78_4 | spl78_163), inference(avatar_split_clause, [], [f4044, f1738, f702, f698])).
fof(f698, plain, (spl78_3 <=> (e2 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl78_3])])).
fof(f1738, plain, (spl78_163 <=> (e2 = op(e3, op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl78_163])])).
fof(f4044, plain, (~ (e2 = op(e3, e3)) | (~ spl78_4 | spl78_163)), inference(backward_demodulation, [], [f1740, f704])).
fof(f1740, plain, (~ (e2 = op(e3, op(e3, e3))) | spl78_163), inference(avatar_component_clause, [], [f1738])).
fof(f4039, plain, (~ spl78_3 | ~ spl78_7), inference(avatar_split_clause, [], [f4037, f715, f698])).
fof(f4037, plain, (~ (e2 = op(e3, e3)) | ~ spl78_7), inference(backward_demodulation, [], [f306, f717])).
fof(f717, plain, ((e2 = op(e3, e2)) | ~ spl78_7), inference(avatar_component_clause, [], [f715])).
fof(f306, plain, ~ (op(e3, e2) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f4025, plain, (~ spl78_3 | ~ spl78_19), inference(avatar_split_clause, [], [f4023, f766, f698])).
fof(f4023, plain, (~ (e2 = op(e3, e3)) | ~ spl78_19), inference(backward_demodulation, [], [f282, f768])).
fof(f768, plain, ((e2 = op(e2, e3)) | ~ spl78_19), inference(avatar_component_clause, [], [f766])).
fof(f282, plain, ~ (op(e2, e3) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f3994, plain, (~ spl78_36 | ~ spl78_44), inference(avatar_split_clause, [], [f3989, f872, f838])).
fof(f838, plain, (spl78_36 <=> (e3 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl78_36])])).
fof(f3989, plain, (~ (e3 = op(e1, e3)) | ~ spl78_44), inference(backward_demodulation, [], [f293, f874])).
fof(f293, plain, ~ (op(e1, e1) = op(e1, e3)), inference(cnf_transformation, [], [f3])).
fof(f3993, plain, (~ spl78_12 | ~ spl78_44), inference(avatar_split_clause, [], [f3987, f872, f736])).
fof(f736, plain, (spl78_12 <=> (e3 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl78_12])])).
fof(f3987, plain, (~ (e3 = op(e3, e1)) | ~ spl78_44), inference(backward_demodulation, [], [f269, f874])).
fof(f269, plain, ~ (op(e1, e1) = op(e3, e1)), inference(cnf_transformation, [], [f3])).
fof(f3961, plain, (~ spl78_6 | spl78_148 | ~ spl78_166), inference(avatar_split_clause, [], [f3960, f1756, f1665, f711])).
fof(f1665, plain, (spl78_148 <=> (e1 = op(op(e2, op(e2, e2)), e2))), introduced(avatar_definition, [new_symbols(naming, [spl78_148])])).
fof(f1756, plain, (spl78_166 <=> (e3 = op(e2, op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl78_166])])).
fof(f3960, plain, (~ (e1 = op(e3, e2)) | (spl78_148 | ~ spl78_166)), inference(forward_demodulation, [], [f1667, f1757])).
fof(f1757, plain, ((e3 = op(e2, op(e2, e2))) | ~ spl78_166), inference(avatar_component_clause, [], [f1756])).
fof(f1667, plain, (~ (e1 = op(op(e2, op(e2, e2)), e2)) | spl78_148), inference(avatar_component_clause, [], [f1665])).
fof(f3953, plain, (~ spl78_52 | ~ spl78_56), inference(avatar_split_clause, [], [f3952, f923, f906])).
fof(f3952, plain, (~ (e3 = op(e0, e3)) | ~ spl78_56), inference(forward_demodulation, [], [f288, f925])).
fof(f288, plain, ~ (op(e0, e2) = op(e0, e3)), inference(cnf_transformation, [], [f3])).
fof(f3951, plain, (~ spl78_51 | ~ spl78_59), inference(avatar_split_clause, [], [f3950, f936, f902])).
fof(f3950, plain, (~ (e2 = op(e0, e3)) | ~ spl78_59), inference(forward_demodulation, [], [f287, f938])).
fof(f938, plain, ((e2 = op(e0, e1)) | ~ spl78_59), inference(avatar_component_clause, [], [f936])).
fof(f3949, plain, (~ spl78_43 | ~ spl78_47), inference(avatar_split_clause, [], [f3948, f885, f868])).
fof(f868, plain, (spl78_43 <=> (e2 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl78_43])])).
fof(f885, plain, (spl78_47 <=> (e2 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl78_47])])).
fof(f3948, plain, (~ (e2 = op(e1, e1)) | ~ spl78_47), inference(forward_demodulation, [], [f289, f887])).
fof(f887, plain, ((e2 = op(e1, e0)) | ~ spl78_47), inference(avatar_component_clause, [], [f885])).
fof(f3947, plain, (~ spl78_43 | ~ spl78_59), inference(avatar_split_clause, [], [f3946, f936, f868])).
fof(f3946, plain, (~ (e2 = op(e1, e1)) | ~ spl78_59), inference(forward_demodulation, [], [f265, f938])).
fof(f3945, plain, (~ spl78_39 | ~ spl78_47), inference(avatar_split_clause, [], [f3944, f885, f851])).
fof(f3944, plain, (~ (e2 = op(e1, e2)) | ~ spl78_47), inference(forward_demodulation, [], [f290, f887])).
fof(f290, plain, ~ (op(e1, e0) = op(e1, e2)), inference(cnf_transformation, [], [f3])).
fof(f3943, plain, (~ spl78_40 | ~ spl78_56), inference(avatar_split_clause, [], [f3942, f923, f855])).
fof(f855, plain, (spl78_40 <=> (e3 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl78_40])])).
fof(f3942, plain, (~ (e3 = op(e1, e2)) | ~ spl78_56), inference(forward_demodulation, [], [f271, f925])).
fof(f271, plain, ~ (op(e0, e2) = op(e1, e2)), inference(cnf_transformation, [], [f3])).
fof(f3941, plain, (~ spl78_28 | ~ spl78_32), inference(avatar_split_clause, [], [f3940, f821, f804])).
fof(f804, plain, (spl78_28 <=> (e3 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl78_28])])).
fof(f3940, plain, (~ (e3 = op(e2, e1)) | ~ spl78_32), inference(forward_demodulation, [], [f295, f823])).
fof(f3939, plain, (~ spl78_27 | ~ spl78_59), inference(avatar_split_clause, [], [f3938, f936, f800])).
fof(f3938, plain, (~ (e2 = op(e2, e1)) | ~ spl78_59), inference(forward_demodulation, [], [f266, f938])).
fof(f3933, plain, (~ spl78_16 | ~ spl78_32), inference(avatar_split_clause, [], [f3932, f821, f753])).
fof(f3932, plain, (~ (e3 = op(e3, e0)) | ~ spl78_32), inference(forward_demodulation, [], [f264, f823])).
fof(f264, plain, ~ (op(e2, e0) = op(e3, e0)), inference(cnf_transformation, [], [f3])).
fof(f3931, plain, (~ spl78_15 | ~ spl78_47), inference(avatar_split_clause, [], [f3930, f885, f749])).
fof(f3930, plain, (~ (e2 = op(e3, e0)) | ~ spl78_47), inference(forward_demodulation, [], [f263, f887])).
fof(f263, plain, ~ (op(e1, e0) = op(e3, e0)), inference(cnf_transformation, [], [f3])).
fof(f3926, plain, (~ spl78_2 | ~ spl78_6), inference(avatar_split_clause, [], [f3923, f711, f694])).
fof(f3923, plain, (~ (e1 = op(e3, e3)) | ~ spl78_6), inference(backward_demodulation, [], [f306, f713])).
fof(f713, plain, ((e1 = op(e3, e2)) | ~ spl78_6), inference(avatar_component_clause, [], [f711])).
fof(f3922, plain, (~ spl78_8 | ~ spl78_12), inference(avatar_split_clause, [], [f3920, f736, f719])).
fof(f719, plain, (spl78_8 <=> (e3 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl78_8])])).
fof(f3920, plain, (~ (e3 = op(e3, e2)) | ~ spl78_12), inference(backward_demodulation, [], [f304, f738])).
fof(f738, plain, ((e3 = op(e3, e1)) | ~ spl78_12), inference(avatar_component_clause, [], [f736])).
fof(f3919, plain, (~ spl78_2 | ~ spl78_18), inference(avatar_split_clause, [], [f3916, f762, f694])).
fof(f3916, plain, (~ (e1 = op(e3, e3)) | ~ spl78_18), inference(backward_demodulation, [], [f282, f764])).
fof(f764, plain, ((e1 = op(e2, e3)) | ~ spl78_18), inference(avatar_component_clause, [], [f762])).
fof(f3915, plain, (~ spl78_20 | ~ spl78_32), inference(avatar_split_clause, [], [f3914, f821, f770])).
fof(f770, plain, (spl78_20 <=> (e3 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl78_20])])).
fof(f3914, plain, (~ (e3 = op(e2, e3)) | ~ spl78_32), inference(backward_demodulation, [], [f297, f823])).
fof(f297, plain, ~ (op(e2, e0) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f3913, plain, (~ spl78_20 | ~ spl78_36), inference(avatar_split_clause, [], [f3911, f838, f770])).
fof(f3911, plain, (~ (e3 = op(e2, e3)) | ~ spl78_36), inference(backward_demodulation, [], [f280, f840])).
fof(f840, plain, ((e3 = op(e1, e3)) | ~ spl78_36), inference(avatar_component_clause, [], [f838])).
fof(f3910, plain, (~ spl78_35 | ~ spl78_47), inference(avatar_split_clause, [], [f3908, f885, f834])).
fof(f3908, plain, (~ (e2 = op(e1, e3)) | ~ spl78_47), inference(backward_demodulation, [], [f291, f887])).
fof(f291, plain, ~ (op(e1, e0) = op(e1, e3)), inference(cnf_transformation, [], [f3])).
fof(f3906, plain, (~ spl78_8 | ~ spl78_56), inference(avatar_split_clause, [], [f3905, f923, f719])).
fof(f3905, plain, (~ (e3 = op(e3, e2)) | ~ spl78_56), inference(backward_demodulation, [], [f273, f925])).
fof(f3904, plain, (~ spl78_11 | ~ spl78_59), inference(avatar_split_clause, [], [f3901, f936, f732])).
fof(f3901, plain, (~ (e2 = op(e3, e1)) | ~ spl78_59), inference(backward_demodulation, [], [f267, f938])).
fof(f267, plain, ~ (op(e0, e1) = op(e3, e1)), inference(cnf_transformation, [], [f3])).
fof(f3899, plain, (~ spl78_54 | ~ spl78_62), inference(avatar_split_clause, [], [f3894, f949, f915])).
fof(f915, plain, (spl78_54 <=> (e1 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl78_54])])).
fof(f3894, plain, (~ (e1 = op(e0, e2)) | ~ spl78_62), inference(backward_demodulation, [], [f284, f951])).
fof(f3898, plain, (~ spl78_30 | ~ spl78_62), inference(avatar_split_clause, [], [f3892, f949, f813])).
fof(f3892, plain, (~ (e1 = op(e2, e0)) | ~ spl78_62), inference(backward_demodulation, [], [f260, f951])).
fof(f3889, plain, (~ spl78_35 | spl78_150 | ~ spl78_157), inference(avatar_split_clause, [], [f3887, f1708, f1674, f834])).
fof(f1674, plain, (spl78_150 <=> (e2 = op(op(e3, op(e3, e3)), e3))), introduced(avatar_definition, [new_symbols(naming, [spl78_150])])).
fof(f1708, plain, (spl78_157 <=> (e1 = op(e3, op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl78_157])])).
fof(f3887, plain, (~ (e2 = op(e1, e3)) | (spl78_150 | ~ spl78_157)), inference(backward_demodulation, [], [f1676, f1709])).
fof(f1709, plain, ((e1 = op(e3, op(e3, e3))) | ~ spl78_157), inference(avatar_component_clause, [], [f1708])).
fof(f1676, plain, (~ (e2 = op(op(e3, op(e3, e3)), e3)) | spl78_150), inference(avatar_component_clause, [], [f1674])).
fof(f3885, plain, (~ spl78_32 | spl78_153 | ~ spl78_159), inference(avatar_split_clause, [], [f3881, f1718, f1688, f821])).
fof(f1688, plain, (spl78_153 <=> (e3 = op(op(e0, op(e0, e0)), e0))), introduced(avatar_definition, [new_symbols(naming, [spl78_153])])).
fof(f3881, plain, (~ (e3 = op(e2, e0)) | (spl78_153 | ~ spl78_159)), inference(backward_demodulation, [], [f1690, f1719])).
fof(f1719, plain, ((e2 = op(e0, op(e0, e0))) | ~ spl78_159), inference(avatar_component_clause, [], [f1718])).
fof(f1690, plain, (~ (e3 = op(op(e0, op(e0, e0)), e0)) | spl78_153), inference(avatar_component_clause, [], [f1688])).
fof(f3863, plain, (~ spl78_3 | ~ spl78_11), inference(avatar_split_clause, [], [f3861, f732, f698])).
fof(f3861, plain, (~ (e2 = op(e3, e3)) | ~ spl78_11), inference(backward_demodulation, [], [f305, f734])).
fof(f305, plain, ~ (op(e3, e1) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f3845, plain, (~ spl78_18 | ~ spl78_30), inference(avatar_split_clause, [], [f3842, f813, f762])).
fof(f3842, plain, (~ (e1 = op(e2, e3)) | ~ spl78_30), inference(backward_demodulation, [], [f297, f815])).
fof(f3839, plain, (~ spl78_3 | ~ spl78_35), inference(avatar_split_clause, [], [f3837, f834, f698])).
fof(f3837, plain, (~ (e2 = op(e3, e3)) | ~ spl78_35), inference(backward_demodulation, [], [f281, f836])).
fof(f281, plain, ~ (op(e1, e3) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f3829, plain, (~ spl78_40 | ~ spl78_48), inference(avatar_split_clause, [], [f3826, f889, f855])).
fof(f3826, plain, (~ (e3 = op(e1, e2)) | ~ spl78_48), inference(backward_demodulation, [], [f290, f891])).
fof(f3828, plain, (~ spl78_16 | ~ spl78_48), inference(avatar_split_clause, [], [f3825, f889, f753])).
fof(f3825, plain, (~ (e3 = op(e3, e0)) | ~ spl78_48), inference(backward_demodulation, [], [f263, f891])).
fof(f3815, plain, (~ spl78_48 | ~ spl78_54 | ~ spl78_63 | spl78_153), inference(avatar_split_clause, [], [f3812, f1688, f953, f915, f889])).
fof(f3812, plain, (~ (e3 = op(e1, e0)) | (~ spl78_54 | ~ spl78_63 | spl78_153)), inference(backward_demodulation, [], [f3795, f917])).
fof(f917, plain, ((e1 = op(e0, e2)) | ~ spl78_54), inference(avatar_component_clause, [], [f915])).
fof(f3795, plain, (~ (e3 = op(op(e0, e2), e0)) | (~ spl78_63 | spl78_153)), inference(backward_demodulation, [], [f1690, f955])).
fof(f3814, plain, (~ spl78_6 | ~ spl78_54), inference(avatar_split_clause, [], [f3808, f915, f711])).
fof(f3808, plain, (~ (e1 = op(e3, e2)) | ~ spl78_54), inference(backward_demodulation, [], [f273, f917])).
fof(f3806, plain, (~ spl78_52 | ~ spl78_60), inference(avatar_split_clause, [], [f3804, f940, f906])).
fof(f3804, plain, (~ (e3 = op(e0, e3)) | ~ spl78_60), inference(backward_demodulation, [], [f287, f942])).
fof(f3800, plain, (~ spl78_54 | ~ spl78_63 | spl78_152), inference(avatar_split_clause, [], [f3794, f1684, f953, f915])).
fof(f1684, plain, (spl78_152 <=> (e1 = op(e0, op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl78_152])])).
fof(f3794, plain, (~ (e1 = op(e0, e2)) | (~ spl78_63 | spl78_152)), inference(backward_demodulation, [], [f1686, f955])).
fof(f1686, plain, (~ (e1 = op(e0, op(e0, e0))) | spl78_152), inference(avatar_component_clause, [], [f1684])).
fof(f3799, plain, (~ spl78_59 | ~ spl78_63), inference(avatar_split_clause, [], [f3790, f953, f936])).
fof(f3790, plain, (~ (e2 = op(e0, e1)) | ~ spl78_63), inference(backward_demodulation, [], [f283, f955])).
fof(f283, plain, ~ (op(e0, e0) = op(e0, e1)), inference(cnf_transformation, [], [f3])).
fof(f3798, plain, (~ spl78_47 | ~ spl78_63), inference(avatar_split_clause, [], [f3787, f953, f885])).
fof(f3787, plain, (~ (e2 = op(e1, e0)) | ~ spl78_63), inference(backward_demodulation, [], [f259, f955])).
fof(f259, plain, ~ (op(e0, e0) = op(e1, e0)), inference(cnf_transformation, [], [f3])).
fof(f3783, plain, (~ spl78_18 | spl78_151 | ~ spl78_163), inference(avatar_split_clause, [], [f3782, f1738, f1679, f762])).
fof(f1679, plain, (spl78_151 <=> (e1 = op(op(e3, op(e3, e3)), e3))), introduced(avatar_definition, [new_symbols(naming, [spl78_151])])).
fof(f3782, plain, (~ (e1 = op(e2, e3)) | (spl78_151 | ~ spl78_163)), inference(backward_demodulation, [], [f1681, f1739])).
fof(f1739, plain, ((e2 = op(e3, op(e3, e3))) | ~ spl78_163), inference(avatar_component_clause, [], [f1738])).
fof(f1681, plain, (~ (e1 = op(op(e3, op(e3, e3)), e3)) | spl78_151), inference(avatar_component_clause, [], [f1679])).
fof(f3759, plain, (~ spl78_17 | ~ spl78_29), inference(avatar_split_clause, [], [f3758, f809, f758])).
fof(f758, plain, (spl78_17 <=> (e0 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl78_17])])).
fof(f3758, plain, (~ (e0 = op(e2, e3)) | ~ spl78_29), inference(backward_demodulation, [], [f297, f811])).
fof(f3754, plain, (~ spl78_17 | ~ spl78_33), inference(avatar_split_clause, [], [f3751, f826, f758])).
fof(f3751, plain, (~ (e0 = op(e2, e3)) | ~ spl78_33), inference(backward_demodulation, [], [f280, f828])).
fof(f828, plain, ((e0 = op(e1, e3)) | ~ spl78_33), inference(avatar_component_clause, [], [f826])).
fof(f3749, plain, (~ spl78_41 | ~ spl78_42 | spl78_143), inference(avatar_split_clause, [], [f3743, f1642, f864, f860])).
fof(f3743, plain, (~ (e0 = op(e1, e1)) | (~ spl78_42 | spl78_143)), inference(backward_demodulation, [], [f1644, f866])).
fof(f3745, plain, (~ spl78_34 | ~ spl78_42), inference(avatar_split_clause, [], [f3739, f864, f830])).
fof(f830, plain, (spl78_34 <=> (e1 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl78_34])])).
fof(f3739, plain, (~ (e1 = op(e1, e3)) | ~ spl78_42), inference(backward_demodulation, [], [f293, f866])).
fof(f3744, plain, (~ spl78_10 | ~ spl78_42), inference(avatar_split_clause, [], [f3738, f864, f728])).
fof(f728, plain, (spl78_10 <=> (e1 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl78_10])])).
fof(f3738, plain, (~ (e1 = op(e3, e1)) | ~ spl78_42), inference(backward_demodulation, [], [f269, f866])).
fof(f3716, plain, (~ spl78_6 | ~ spl78_3 | spl78_157), inference(avatar_split_clause, [], [f3715, f1708, f698, f711])).
fof(f3715, plain, (~ (e1 = op(e3, e2)) | (~ spl78_3 | spl78_157)), inference(forward_demodulation, [], [f1710, f700])).
fof(f700, plain, ((e2 = op(e3, e3)) | ~ spl78_3), inference(avatar_component_clause, [], [f698])).
fof(f1710, plain, (~ (e1 = op(e3, op(e3, e3))) | spl78_157), inference(avatar_component_clause, [], [f1708])).
fof(f3702, plain, (~ spl78_64 | ~ spl78_52), inference(avatar_split_clause, [], [f3597, f906, f957])).
fof(f957, plain, (spl78_64 <=> (op(e0, e0) = e3)), introduced(avatar_definition, [new_symbols(naming, [spl78_64])])).
fof(f3597, plain, (~ (op(e0, e0) = e3) | ~ spl78_52), inference(forward_demodulation, [], [f285, f908])).
fof(f908, plain, ((e3 = op(e0, e3)) | ~ spl78_52), inference(avatar_component_clause, [], [f906])).
fof(f3700, plain, (~ spl78_64 | ~ spl78_16), inference(avatar_split_clause, [], [f3699, f753, f957])).
fof(f3699, plain, (~ (op(e0, e0) = e3) | ~ spl78_16), inference(forward_demodulation, [], [f261, f755])).
fof(f755, plain, ((e3 = op(e3, e0)) | ~ spl78_16), inference(avatar_component_clause, [], [f753])).
fof(f3690, plain, (~ spl78_28 | spl78_144 | ~ spl78_161), inference(avatar_contradiction_clause, [], [f3689])).
fof(f3689, plain, ($false | (~ spl78_28 | spl78_144 | ~ spl78_161)), inference(subsumption_resolution, [], [f3688, f806])).
fof(f806, plain, ((e3 = op(e2, e1)) | ~ spl78_28), inference(avatar_component_clause, [], [f804])).
fof(f3688, plain, (~ (e3 = op(e2, e1)) | (spl78_144 | ~ spl78_161)), inference(forward_demodulation, [], [f1648, f1729])).
fof(f1648, plain, (~ (e3 = op(op(e1, op(e1, e1)), e1)) | spl78_144), inference(avatar_component_clause, [], [f1646])).
fof(f1646, plain, (spl78_144 <=> (e3 = op(op(e1, op(e1, e1)), e1))), introduced(avatar_definition, [new_symbols(naming, [spl78_144])])).
fof(f3680, plain, (spl78_5 | ~ spl78_3 | ~ spl78_149), inference(avatar_split_clause, [], [f3679, f1670, f698, f707])).
fof(f1670, plain, (spl78_149 <=> (e0 = op(e3, op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl78_149])])).
fof(f3679, plain, ((e0 = op(e3, e2)) | (~ spl78_3 | ~ spl78_149)), inference(backward_demodulation, [], [f1671, f700])).
fof(f1671, plain, ((e0 = op(e3, op(e3, e3))) | ~ spl78_149), inference(avatar_component_clause, [], [f1670])).
fof(f3677, plain, (~ spl78_1 | ~ spl78_5), inference(avatar_split_clause, [], [f3673, f707, f690])).
fof(f690, plain, (spl78_1 <=> (e0 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl78_1])])).
fof(f3673, plain, (~ (e0 = op(e3, e3)) | ~ spl78_5), inference(backward_demodulation, [], [f306, f709])).
fof(f709, plain, ((e0 = op(e3, e2)) | ~ spl78_5), inference(avatar_component_clause, [], [f707])).
fof(f3668, plain, (~ spl78_8 | ~ spl78_16), inference(avatar_split_clause, [], [f3666, f753, f719])).
fof(f3666, plain, (~ (e3 = op(e3, e2)) | ~ spl78_16), inference(backward_demodulation, [], [f302, f755])).
fof(f3664, plain, (~ spl78_1 | ~ spl78_17), inference(avatar_split_clause, [], [f3661, f758, f690])).
fof(f3661, plain, (~ (e0 = op(e3, e3)) | ~ spl78_17), inference(backward_demodulation, [], [f282, f760])).
fof(f760, plain, ((e0 = op(e2, e3)) | ~ spl78_17), inference(avatar_component_clause, [], [f758])).
fof(f3659, plain, (spl78_21 | ~ spl78_23 | ~ spl78_156), inference(avatar_split_clause, [], [f3658, f1703, f783, f775])).
fof(f1703, plain, (spl78_156 <=> (e0 = op(op(e2, op(e2, e2)), e2))), introduced(avatar_definition, [new_symbols(naming, [spl78_156])])).
fof(f3658, plain, ((e0 = op(e2, e2)) | (~ spl78_23 | ~ spl78_156)), inference(forward_demodulation, [], [f3656, f785])).
fof(f785, plain, ((e2 = op(e2, e2)) | ~ spl78_23), inference(avatar_component_clause, [], [f783])).
fof(f3656, plain, ((e0 = op(op(e2, e2), e2)) | (~ spl78_23 | ~ spl78_156)), inference(backward_demodulation, [], [f1704, f785])).
fof(f1704, plain, ((e0 = op(op(e2, op(e2, e2)), e2)) | ~ spl78_156), inference(avatar_component_clause, [], [f1703])).
fof(f3657, plain, (~ spl78_21 | ~ spl78_23 | spl78_146), inference(avatar_split_clause, [], [f3654, f1656, f783, f775])).
fof(f3654, plain, (~ (e0 = op(e2, e2)) | (~ spl78_23 | spl78_146)), inference(backward_demodulation, [], [f1658, f785])).
fof(f1658, plain, (~ (e0 = op(e2, op(e2, e2))) | spl78_146), inference(avatar_component_clause, [], [f1656])).
fof(f3643, plain, (~ spl78_18 | ~ spl78_34), inference(avatar_split_clause, [], [f3639, f830, f762])).
fof(f3639, plain, (~ (e1 = op(e2, e3)) | ~ spl78_34), inference(backward_demodulation, [], [f280, f832])).
fof(f832, plain, ((e1 = op(e1, e3)) | ~ spl78_34), inference(avatar_component_clause, [], [f830])).
fof(f3638, plain, (~ spl78_8 | ~ spl78_40), inference(avatar_split_clause, [], [f3636, f855, f719])).
fof(f3636, plain, (~ (e3 = op(e3, e2)) | ~ spl78_40), inference(backward_demodulation, [], [f275, f857])).
fof(f857, plain, ((e3 = op(e1, e2)) | ~ spl78_40), inference(avatar_component_clause, [], [f855])).
fof(f3631, plain, (~ spl78_31 | ~ spl78_47), inference(avatar_split_clause, [], [f3624, f885, f817])).
fof(f3624, plain, (~ (e2 = op(e2, e0)) | ~ spl78_47), inference(backward_demodulation, [], [f262, f887])).
fof(f3623, plain, (~ spl78_38 | ~ spl78_54), inference(avatar_split_clause, [], [f3618, f915, f847])).
fof(f3618, plain, (~ (e1 = op(e1, e2)) | ~ spl78_54), inference(backward_demodulation, [], [f271, f917])).
fof(f3617, plain, (~ spl78_55 | ~ spl78_59), inference(avatar_split_clause, [], [f3615, f936, f919])).
fof(f3615, plain, (~ (e2 = op(e0, e2)) | ~ spl78_59), inference(backward_demodulation, [], [f286, f938])).
fof(f3571, plain, (~ spl78_10 | ~ spl78_14), inference(avatar_split_clause, [], [f3567, f745, f728])).
fof(f3567, plain, (~ (e1 = op(e3, e1)) | ~ spl78_14), inference(backward_demodulation, [], [f301, f747])).
fof(f3559, plain, (~ spl78_17 | ~ spl78_21), inference(avatar_split_clause, [], [f3551, f775, f758])).
fof(f3551, plain, (~ (e0 = op(e2, e3)) | ~ spl78_21), inference(backward_demodulation, [], [f300, f777])).
fof(f777, plain, ((e0 = op(e2, e2)) | ~ spl78_21), inference(avatar_component_clause, [], [f775])).
fof(f300, plain, ~ (op(e2, e2) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f3546, plain, (~ spl78_22 | ~ spl78_38), inference(avatar_split_clause, [], [f3544, f847, f779])).
fof(f779, plain, (spl78_22 <=> (e1 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl78_22])])).
fof(f3544, plain, (~ (e1 = op(e2, e2)) | ~ spl78_38), inference(backward_demodulation, [], [f274, f849])).
fof(f274, plain, ~ (op(e1, e2) = op(e2, e2)), inference(cnf_transformation, [], [f3])).
fof(f3541, plain, (~ spl78_41 | spl78_45 | ~ spl78_143), inference(avatar_contradiction_clause, [], [f3540])).
fof(f3540, plain, ($false | (~ spl78_41 | spl78_45 | ~ spl78_143)), inference(subsumption_resolution, [], [f3538, f878])).
fof(f878, plain, (~ (e0 = op(e1, e0)) | spl78_45), inference(avatar_component_clause, [], [f877])).
fof(f3538, plain, ((e0 = op(e1, e0)) | (~ spl78_41 | ~ spl78_143)), inference(backward_demodulation, [], [f1643, f862])).
fof(f1643, plain, ((e0 = op(e1, op(e1, e1))) | ~ spl78_143), inference(avatar_component_clause, [], [f1642])).
fof(f3531, plain, (~ spl78_36 | ~ spl78_48), inference(avatar_split_clause, [], [f3530, f889, f838])).
fof(f3530, plain, (~ (e3 = op(e1, e3)) | ~ spl78_48), inference(backward_demodulation, [], [f291, f891])).
fof(f3526, plain, (~ spl78_36 | ~ spl78_52), inference(avatar_split_clause, [], [f3523, f906, f838])).
fof(f3523, plain, (~ (e3 = op(e1, e3)) | ~ spl78_52), inference(backward_demodulation, [], [f277, f908])).
fof(f3521, plain, (~ spl78_10 | ~ spl78_58), inference(avatar_split_clause, [], [f3517, f932, f728])).
fof(f3517, plain, (~ (e1 = op(e3, e1)) | ~ spl78_58), inference(backward_demodulation, [], [f267, f934])).
fof(f3515, plain, (~ spl78_61 | ~ spl78_63), inference(avatar_contradiction_clause, [], [f3514])).
fof(f3514, plain, ($false | (~ spl78_61 | ~ spl78_63)), inference(subsumption_resolution, [], [f3513, f308])).
fof(f308, plain, ~ (e0 = e2), inference(cnf_transformation, [], [f4])).
fof(f3513, plain, ((e0 = e2) | (~ spl78_61 | ~ spl78_63)), inference(backward_demodulation, [], [f955, f947])).
fof(f3476, plain, (~ spl78_3 | spl78_7 | ~ spl78_163), inference(avatar_contradiction_clause, [], [f3475])).
fof(f3475, plain, ($false | (~ spl78_3 | spl78_7 | ~ spl78_163)), inference(subsumption_resolution, [], [f3474, f716])).
fof(f716, plain, (~ (e2 = op(e3, e2)) | spl78_7), inference(avatar_component_clause, [], [f715])).
fof(f3474, plain, ((e2 = op(e3, e2)) | (~ spl78_3 | ~ spl78_163)), inference(backward_demodulation, [], [f1739, f700])).
fof(f3469, plain, (~ spl78_6 | ~ spl78_10), inference(avatar_split_clause, [], [f3465, f728, f711])).
fof(f3465, plain, (~ (e1 = op(e3, e2)) | ~ spl78_10), inference(backward_demodulation, [], [f304, f730])).
fof(f730, plain, ((e1 = op(e3, e1)) | ~ spl78_10), inference(avatar_component_clause, [], [f728])).
fof(f3464, plain, (~ spl78_1 | ~ spl78_13), inference(avatar_split_clause, [], [f3460, f741, f690])).
fof(f3460, plain, (~ (e0 = op(e3, e3)) | ~ spl78_13), inference(backward_demodulation, [], [f303, f743])).
fof(f3452, plain, (~ spl78_6 | ~ spl78_22), inference(avatar_split_clause, [], [f3445, f779, f711])).
fof(f3445, plain, (~ (e1 = op(e3, e2)) | ~ spl78_22), inference(backward_demodulation, [], [f276, f781])).
fof(f781, plain, ((e1 = op(e2, e2)) | ~ spl78_22), inference(avatar_component_clause, [], [f779])).
fof(f276, plain, ~ (op(e2, e2) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f3444, plain, (~ spl78_24 | ~ spl78_28), inference(avatar_split_clause, [], [f3441, f804, f787])).
fof(f3441, plain, (~ (e3 = op(e2, e2)) | ~ spl78_28), inference(backward_demodulation, [], [f298, f806])).
fof(f298, plain, ~ (op(e2, e1) = op(e2, e2)), inference(cnf_transformation, [], [f3])).
fof(f3443, plain, (~ spl78_12 | ~ spl78_28), inference(avatar_split_clause, [], [f3440, f804, f736])).
fof(f3440, plain, (~ (e3 = op(e3, e1)) | ~ spl78_28), inference(backward_demodulation, [], [f270, f806])).
fof(f3439, plain, (~ spl78_19 | ~ spl78_31), inference(avatar_split_clause, [], [f3437, f817, f766])).
fof(f3437, plain, (~ (e2 = op(e2, e3)) | ~ spl78_31), inference(backward_demodulation, [], [f297, f819])).
fof(f3438, plain, (~ spl78_15 | ~ spl78_31), inference(avatar_split_clause, [], [f3433, f817, f749])).
fof(f3433, plain, (~ (e2 = op(e3, e0)) | ~ spl78_31), inference(backward_demodulation, [], [f264, f819])).
fof(f3431, plain, (~ spl78_49 | ~ spl78_50), inference(avatar_contradiction_clause, [], [f3430])).
fof(f3430, plain, ($false | (~ spl78_49 | ~ spl78_50)), inference(subsumption_resolution, [], [f3429, f307])).
fof(f307, plain, ~ (e0 = e1), inference(cnf_transformation, [], [f4])).
fof(f3429, plain, ((e0 = e1) | (~ spl78_49 | ~ spl78_50)), inference(backward_demodulation, [], [f900, f896])).
fof(f900, plain, ((e1 = op(e0, e3)) | ~ spl78_50), inference(avatar_component_clause, [], [f898])).
fof(f3423, plain, (~ spl78_49 | ~ spl78_57), inference(avatar_split_clause, [], [f3337, f928, f894])).
fof(f928, plain, (spl78_57 <=> (e0 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl78_57])])).
fof(f3337, plain, (~ (e0 = op(e0, e3)) | ~ spl78_57), inference(backward_demodulation, [], [f287, f930])).
fof(f930, plain, ((e0 = op(e0, e1)) | ~ spl78_57), inference(avatar_component_clause, [], [f928])).
fof(f3362, plain, (~ spl78_38 | ~ spl78_46), inference(avatar_split_clause, [], [f3356, f881, f847])).
fof(f3356, plain, (~ (e1 = op(e1, e2)) | ~ spl78_46), inference(backward_demodulation, [], [f290, f883])).
fof(f3360, plain, (~ spl78_30 | ~ spl78_46), inference(avatar_split_clause, [], [f3353, f881, f813])).
fof(f3353, plain, (~ (e1 = op(e2, e0)) | ~ spl78_46), inference(backward_demodulation, [], [f262, f883])).
fof(f3352, plain, (~ spl78_2 | ~ spl78_50), inference(avatar_split_clause, [], [f3349, f898, f694])).
fof(f3349, plain, (~ (e1 = op(e3, e3)) | ~ spl78_50), inference(backward_demodulation, [], [f279, f900])).
fof(f3347, plain, (~ spl78_7 | ~ spl78_55), inference(avatar_split_clause, [], [f3344, f919, f715])).
fof(f3344, plain, (~ (e2 = op(e3, e2)) | ~ spl78_55), inference(backward_demodulation, [], [f273, f921])).
fof(f3341, plain, (~ spl78_9 | ~ spl78_57), inference(avatar_split_clause, [], [f3335, f928, f724])).
fof(f3335, plain, (~ (e0 = op(e3, e1)) | ~ spl78_57), inference(backward_demodulation, [], [f267, f930])).
fof(f3322, plain, (~ spl78_59 | ~ spl78_143 | spl78_145), inference(avatar_split_clause, [], [f3320, f1651, f1642, f936])).
fof(f3320, plain, (~ (e2 = op(e0, e1)) | (~ spl78_143 | spl78_145)), inference(backward_demodulation, [], [f1653, f1643])).
fof(f3291, plain, (~ spl78_12 | ~ spl78_16), inference(avatar_split_clause, [], [f3288, f753, f736])).
fof(f3288, plain, (~ (e3 = op(e3, e1)) | ~ spl78_16), inference(backward_demodulation, [], [f301, f755])).
fof(f3283, plain, (~ spl78_5 | ~ spl78_21), inference(avatar_split_clause, [], [f3276, f775, f707])).
fof(f3276, plain, (~ (e0 = op(e3, e2)) | ~ spl78_21), inference(backward_demodulation, [], [f276, f777])).
fof(f3257, plain, (~ spl78_36 | ~ spl78_40), inference(avatar_split_clause, [], [f3256, f855, f838])).
fof(f3256, plain, (~ (e3 = op(e1, e3)) | ~ spl78_40), inference(backward_demodulation, [], [f294, f857])).
fof(f3233, plain, (~ spl78_50 | ~ spl78_54), inference(avatar_split_clause, [], [f3229, f915, f898])).
fof(f3229, plain, (~ (e1 = op(e0, e3)) | ~ spl78_54), inference(backward_demodulation, [], [f288, f917])).
fof(f3223, plain, (~ spl78_40 | spl78_147 | ~ spl78_155), inference(avatar_split_clause, [], [f3219, f1698, f1660, f855])).
fof(f3219, plain, (~ (e3 = op(e1, e2)) | (spl78_147 | ~ spl78_155)), inference(backward_demodulation, [], [f1662, f1699])).
fof(f1699, plain, ((e1 = op(e2, op(e2, e2))) | ~ spl78_155), inference(avatar_component_clause, [], [f1698])).
fof(f3202, plain, (~ spl78_41 | ~ spl78_47 | spl78_161), inference(avatar_contradiction_clause, [], [f3201])).
fof(f3201, plain, ($false | (~ spl78_41 | ~ spl78_47 | spl78_161)), inference(subsumption_resolution, [], [f3195, f887])).
fof(f3195, plain, (~ (e2 = op(e1, e0)) | (~ spl78_41 | spl78_161)), inference(backward_demodulation, [], [f1730, f862])).
fof(f1730, plain, (~ (e2 = op(e1, op(e1, e1))) | spl78_161), inference(avatar_component_clause, [], [f1728])).
fof(f3179, plain, (~ spl78_50 | ~ spl78_3 | ~ spl78_5 | spl78_151), inference(avatar_split_clause, [], [f3178, f1679, f707, f698, f898])).
fof(f3178, plain, (~ (e1 = op(e0, e3)) | (~ spl78_3 | ~ spl78_5 | spl78_151)), inference(forward_demodulation, [], [f3177, f709])).
fof(f3177, plain, (~ (e1 = op(op(e3, e2), e3)) | (~ spl78_3 | spl78_151)), inference(forward_demodulation, [], [f1681, f700])).
fof(f3167, plain, (~ spl78_3 | ~ spl78_5 | spl78_149), inference(avatar_contradiction_clause, [], [f3166])).
fof(f3166, plain, ($false | (~ spl78_3 | ~ spl78_5 | spl78_149)), inference(subsumption_resolution, [], [f3163, f709])).
fof(f3163, plain, (~ (e0 = op(e3, e2)) | (~ spl78_3 | spl78_149)), inference(backward_demodulation, [], [f1672, f700])).
fof(f1672, plain, (~ (e0 = op(e3, op(e3, e3))) | spl78_149), inference(avatar_component_clause, [], [f1670])).
fof(f3154, plain, (~ spl78_2 | ~ spl78_14), inference(avatar_split_clause, [], [f3151, f745, f694])).
fof(f3151, plain, (~ (e1 = op(e3, e3)) | ~ spl78_14), inference(backward_demodulation, [], [f303, f747])).
fof(f3144, plain, (~ spl78_17 | ~ spl78_25), inference(avatar_split_clause, [], [f3140, f792, f758])).
fof(f3140, plain, (~ (e0 = op(e2, e3)) | ~ spl78_25), inference(backward_demodulation, [], [f299, f794])).
fof(f3131, plain, (~ spl78_38 | ~ spl78_40), inference(avatar_contradiction_clause, [], [f3130])).
fof(f3130, plain, ($false | (~ spl78_38 | ~ spl78_40)), inference(subsumption_resolution, [], [f3129, f311])).
fof(f3129, plain, ((e1 = e3) | (~ spl78_38 | ~ spl78_40)), inference(forward_demodulation, [], [f857, f849])).
fof(f3079, plain, (~ spl78_1 | ~ spl78_9), inference(avatar_split_clause, [], [f3078, f724, f690])).
fof(f3078, plain, (~ (e0 = op(e3, e3)) | ~ spl78_9), inference(backward_demodulation, [], [f305, f726])).
fof(f3066, plain, (~ spl78_30 | ~ spl78_32), inference(avatar_contradiction_clause, [], [f3065])).
fof(f3065, plain, ($false | (~ spl78_30 | ~ spl78_32)), inference(subsumption_resolution, [], [f3064, f311])).
fof(f3064, plain, ((e1 = e3) | (~ spl78_30 | ~ spl78_32)), inference(forward_demodulation, [], [f823, f815])).
fof(f3058, plain, (~ spl78_34 | ~ spl78_38), inference(avatar_split_clause, [], [f3055, f847, f830])).
fof(f3055, plain, (~ (e1 = op(e1, e3)) | ~ spl78_38), inference(backward_demodulation, [], [f294, f849])).
fof(f3054, plain, (~ spl78_35 | ~ spl78_44 | spl78_161), inference(avatar_split_clause, [], [f3052, f1728, f872, f834])).
fof(f3052, plain, (~ (e2 = op(e1, e3)) | (~ spl78_44 | spl78_161)), inference(backward_demodulation, [], [f1730, f874])).
fof(f3045, plain, (~ spl78_37 | ~ spl78_45), inference(avatar_split_clause, [], [f3040, f877, f843])).
fof(f3040, plain, (~ (e0 = op(e1, e2)) | ~ spl78_45), inference(backward_demodulation, [], [f290, f879])).
fof(f3038, plain, (~ spl78_20 | ~ spl78_52), inference(avatar_split_clause, [], [f3035, f906, f770])).
fof(f3035, plain, (~ (e3 = op(e2, e3)) | ~ spl78_52), inference(backward_demodulation, [], [f278, f908])).
fof(f3033, plain, (~ spl78_53 | ~ spl78_55), inference(avatar_contradiction_clause, [], [f3032])).
fof(f3032, plain, ($false | (~ spl78_53 | ~ spl78_55)), inference(subsumption_resolution, [], [f3031, f308])).
fof(f3031, plain, ((e0 = e2) | (~ spl78_53 | ~ spl78_55)), inference(backward_demodulation, [], [f921, f913])).
fof(f3029, plain, (~ spl78_51 | ~ spl78_55), inference(avatar_split_clause, [], [f3028, f919, f902])).
fof(f3028, plain, (~ (e2 = op(e0, e3)) | ~ spl78_55), inference(backward_demodulation, [], [f288, f921])).
fof(f3025, plain, (~ spl78_54 | ~ spl78_58), inference(avatar_split_clause, [], [f3019, f932, f915])).
fof(f3019, plain, (~ (e1 = op(e0, e2)) | ~ spl78_58), inference(backward_demodulation, [], [f286, f934])).
fof(f3015, plain, (~ spl78_51 | ~ spl78_149 | spl78_150), inference(avatar_split_clause, [], [f3014, f1674, f1670, f902])).
fof(f3014, plain, (~ (e2 = op(e0, e3)) | (~ spl78_149 | spl78_150)), inference(backward_demodulation, [], [f1676, f1671])).
fof(f3008, plain, (~ spl78_53 | ~ spl78_61), inference(avatar_split_clause, [], [f2630, f945, f911])).
fof(f2630, plain, (~ (e0 = op(e0, e2)) | ~ spl78_61), inference(forward_demodulation, [], [f284, f947])).
fof(f3006, plain, (~ spl78_45 | ~ spl78_61), inference(avatar_split_clause, [], [f2550, f945, f877])).
fof(f2550, plain, (~ (e0 = op(e1, e0)) | ~ spl78_61), inference(backward_demodulation, [], [f259, f947])).
fof(f3003, plain, (~ spl78_27 | ~ spl78_23), inference(avatar_split_clause, [], [f3002, f783, f800])).
fof(f3002, plain, (~ (e2 = op(e2, e1)) | ~ spl78_23), inference(forward_demodulation, [], [f298, f785])).
fof(f2999, plain, (~ spl78_4 | ~ spl78_8), inference(avatar_split_clause, [], [f2998, f719, f702])).
fof(f2998, plain, (~ (e3 = op(e3, e3)) | ~ spl78_8), inference(forward_demodulation, [], [f306, f721])).
fof(f721, plain, ((e3 = op(e3, e2)) | ~ spl78_8), inference(avatar_component_clause, [], [f719])).
fof(f2997, plain, (~ spl78_3 | ~ spl78_15), inference(avatar_split_clause, [], [f2996, f749, f698])).
fof(f2996, plain, (~ (e2 = op(e3, e3)) | ~ spl78_15), inference(forward_demodulation, [], [f303, f751])).
fof(f2976, plain, (~ spl78_21 | ~ spl78_25), inference(avatar_split_clause, [], [f2972, f792, f775])).
fof(f2972, plain, (~ (e0 = op(e2, e2)) | ~ spl78_25), inference(backward_demodulation, [], [f298, f794])).
fof(f2959, plain, (~ spl78_21 | ~ spl78_37), inference(avatar_split_clause, [], [f2953, f843, f775])).
fof(f2953, plain, (~ (e0 = op(e2, e2)) | ~ spl78_37), inference(backward_demodulation, [], [f274, f845])).
fof(f2951, plain, (~ spl78_39 | ~ spl78_43), inference(avatar_split_clause, [], [f2946, f868, f851])).
fof(f2946, plain, (~ (e2 = op(e1, e2)) | ~ spl78_43), inference(backward_demodulation, [], [f292, f870])).
fof(f870, plain, ((e2 = op(e1, e1)) | ~ spl78_43), inference(avatar_component_clause, [], [f868])).
fof(f2950, plain, (~ spl78_27 | ~ spl78_43), inference(avatar_split_clause, [], [f2944, f868, f800])).
fof(f2944, plain, (~ (e2 = op(e2, e1)) | ~ spl78_43), inference(backward_demodulation, [], [f268, f870])).
fof(f2930, plain, (~ spl78_12 | ~ spl78_60), inference(avatar_split_clause, [], [f2926, f940, f736])).
fof(f2926, plain, (~ (e3 = op(e3, e1)) | ~ spl78_60), inference(backward_demodulation, [], [f267, f942])).
fof(f2925, plain, (~ spl78_60 | ~ spl78_143 | spl78_144), inference(avatar_split_clause, [], [f2922, f1646, f1642, f940])).
fof(f2922, plain, (~ (e3 = op(e0, e1)) | (~ spl78_143 | spl78_144)), inference(backward_demodulation, [], [f1648, f1643])).
fof(f2912, plain, (~ spl78_21 | ~ spl78_24), inference(avatar_contradiction_clause, [], [f2911])).
fof(f2911, plain, ($false | (~ spl78_21 | ~ spl78_24)), inference(subsumption_resolution, [], [f2910, f309])).
fof(f309, plain, ~ (e0 = e3), inference(cnf_transformation, [], [f4])).
fof(f2910, plain, ((e0 = e3) | (~ spl78_21 | ~ spl78_24)), inference(backward_demodulation, [], [f789, f777])).
fof(f2905, plain, (~ spl78_8 | ~ spl78_24), inference(avatar_split_clause, [], [f2901, f787, f719])).
fof(f2901, plain, (~ (e3 = op(e3, e2)) | ~ spl78_24), inference(backward_demodulation, [], [f276, f789])).
fof(f2877, plain, (~ spl78_9 | ~ spl78_41), inference(avatar_split_clause, [], [f2867, f860, f724])).
fof(f2867, plain, (~ (e0 = op(e3, e1)) | ~ spl78_41), inference(backward_demodulation, [], [f269, f862])).
fof(f2865, plain, (~ spl78_34 | ~ spl78_46), inference(avatar_split_clause, [], [f2861, f881, f830])).
fof(f2861, plain, (~ (e1 = op(e1, e3)) | ~ spl78_46), inference(backward_demodulation, [], [f291, f883])).
fof(f2857, plain, (~ spl78_49 | ~ spl78_51), inference(avatar_contradiction_clause, [], [f2856])).
fof(f2856, plain, ($false | (~ spl78_49 | ~ spl78_51)), inference(subsumption_resolution, [], [f2855, f308])).
fof(f2855, plain, ((e0 = e2) | (~ spl78_49 | ~ spl78_51)), inference(backward_demodulation, [], [f904, f896])).
fof(f2839, plain, (~ spl78_5 | spl78_156 | ~ spl78_166), inference(avatar_split_clause, [], [f2836, f1756, f1703, f707])).
fof(f2836, plain, (~ (e0 = op(e3, e2)) | (spl78_156 | ~ spl78_166)), inference(backward_demodulation, [], [f1705, f1757])).
fof(f1705, plain, (~ (e0 = op(op(e2, op(e2, e2)), e2)) | spl78_156), inference(avatar_component_clause, [], [f1703])).
fof(f2833, plain, (~ spl78_49 | ~ spl78_1), inference(avatar_split_clause, [], [f2745, f690, f894])).
fof(f2745, plain, (~ (e0 = op(e0, e3)) | ~ spl78_1), inference(forward_demodulation, [], [f279, f692])).
fof(f692, plain, ((e0 = op(e3, e3)) | ~ spl78_1), inference(avatar_component_clause, [], [f690])).
fof(f2823, plain, (~ spl78_18 | ~ spl78_1 | ~ spl78_15 | spl78_151), inference(avatar_split_clause, [], [f2822, f1679, f749, f690, f762])).
fof(f2822, plain, (~ (e1 = op(e2, e3)) | (~ spl78_1 | ~ spl78_15 | spl78_151)), inference(forward_demodulation, [], [f2761, f751])).
fof(f2761, plain, (~ (e1 = op(op(e3, e0), e3)) | (~ spl78_1 | spl78_151)), inference(forward_demodulation, [], [f1681, f692])).
fof(f2816, plain, (~ spl78_14 | ~ spl78_15), inference(avatar_contradiction_clause, [], [f2815])).
fof(f2815, plain, ($false | (~ spl78_14 | ~ spl78_15)), inference(subsumption_resolution, [], [f2814, f310])).
fof(f310, plain, ~ (e1 = e2), inference(cnf_transformation, [], [f4])).
fof(f2814, plain, ((e1 = e2) | (~ spl78_14 | ~ spl78_15)), inference(forward_demodulation, [], [f751, f747])).
fof(f2812, plain, (~ spl78_25 | ~ spl78_28), inference(avatar_contradiction_clause, [], [f2811])).
fof(f2811, plain, ($false | (~ spl78_25 | ~ spl78_28)), inference(subsumption_resolution, [], [f2810, f309])).
fof(f2810, plain, ((e0 = e3) | (~ spl78_25 | ~ spl78_28)), inference(backward_demodulation, [], [f806, f794])).
fof(f2805, plain, (~ spl78_20 | ~ spl78_28), inference(avatar_split_clause, [], [f2800, f804, f770])).
fof(f2800, plain, (~ (e3 = op(e2, e3)) | ~ spl78_28), inference(backward_demodulation, [], [f299, f806])).
fof(f2794, plain, (~ spl78_5 | ~ spl78_37), inference(avatar_split_clause, [], [f2790, f843, f707])).
fof(f2790, plain, (~ (e0 = op(e3, e2)) | ~ spl78_37), inference(backward_demodulation, [], [f275, f845])).
fof(f2759, plain, (~ spl78_1 | ~ spl78_14 | spl78_157), inference(avatar_contradiction_clause, [], [f2758])).
fof(f2758, plain, ($false | (~ spl78_1 | ~ spl78_14 | spl78_157)), inference(subsumption_resolution, [], [f2757, f747])).
fof(f2757, plain, (~ (e1 = op(e3, e0)) | (~ spl78_1 | spl78_157)), inference(forward_demodulation, [], [f1710, f692])).
fof(f2729, plain, (~ spl78_28 | ~ spl78_22 | spl78_166), inference(avatar_split_clause, [], [f2728, f1756, f779, f804])).
fof(f2728, plain, (~ (e3 = op(e2, e1)) | (~ spl78_22 | spl78_166)), inference(forward_demodulation, [], [f1758, f781])).
fof(f1758, plain, (~ (e3 = op(e2, op(e2, e2))) | spl78_166), inference(avatar_component_clause, [], [f1756])).
fof(f2727, plain, (~ spl78_30 | ~ spl78_22), inference(avatar_split_clause, [], [f2726, f779, f813])).
fof(f2726, plain, (~ (e1 = op(e2, e0)) | ~ spl78_22), inference(forward_demodulation, [], [f296, f781])).
fof(f2699, plain, (~ spl78_21 | ~ spl78_22), inference(avatar_contradiction_clause, [], [f2698])).
fof(f2698, plain, ($false | (~ spl78_21 | ~ spl78_22)), inference(subsumption_resolution, [], [f2697, f307])).
fof(f2697, plain, ((e0 = e1) | (~ spl78_21 | ~ spl78_22)), inference(forward_demodulation, [], [f781, f777])).
fof(f2696, plain, (~ spl78_13 | ~ spl78_29), inference(avatar_split_clause, [], [f2694, f809, f741])).
fof(f2694, plain, (~ (e0 = op(e3, e0)) | ~ spl78_29), inference(backward_demodulation, [], [f264, f811])).
fof(f2688, plain, (~ spl78_2 | ~ spl78_34), inference(avatar_split_clause, [], [f2687, f830, f694])).
fof(f2687, plain, (~ (e1 = op(e3, e3)) | ~ spl78_34), inference(backward_demodulation, [], [f281, f832])).
fof(f2656, plain, (~ spl78_39 | ~ spl78_55), inference(avatar_split_clause, [], [f2653, f919, f851])).
fof(f2653, plain, (~ (e2 = op(e1, e2)) | ~ spl78_55), inference(backward_demodulation, [], [f271, f921])).
fof(f2645, plain, (~ spl78_61 | ~ spl78_64), inference(avatar_contradiction_clause, [], [f2644])).
fof(f2644, plain, ($false | (~ spl78_61 | ~ spl78_64)), inference(subsumption_resolution, [], [f2643, f309])).
fof(f2643, plain, ((e0 = e3) | (~ spl78_61 | ~ spl78_64)), inference(forward_demodulation, [], [f959, f947])).
fof(f959, plain, ((op(e0, e0) = e3) | ~ spl78_64), inference(avatar_component_clause, [], [f957])).
fof(f2639, plain, (~ spl78_61 | ~ spl78_159), inference(avatar_contradiction_clause, [], [f2638])).
fof(f2638, plain, ($false | (~ spl78_61 | ~ spl78_159)), inference(subsumption_resolution, [], [f2637, f308])).
fof(f2637, plain, ((e0 = e2) | (~ spl78_61 | ~ spl78_159)), inference(forward_demodulation, [], [f2636, f947])).
fof(f2636, plain, ((op(e0, e0) = e2) | (~ spl78_61 | ~ spl78_159)), inference(forward_demodulation, [], [f1719, f947])).
fof(f2627, plain, (~ spl78_29 | ~ spl78_61), inference(avatar_split_clause, [], [f2626, f945, f809])).
fof(f2626, plain, (~ (e0 = op(e2, e0)) | ~ spl78_61), inference(forward_demodulation, [], [f260, f947])).
fof(f2621, plain, (~ spl78_32 | ~ spl78_21 | spl78_166), inference(avatar_split_clause, [], [f2545, f1756, f775, f821])).
fof(f2545, plain, (~ (e3 = op(e2, e0)) | (~ spl78_21 | spl78_166)), inference(forward_demodulation, [], [f1758, f777])).
fof(f2616, plain, (~ spl78_4 | ~ spl78_20), inference(avatar_split_clause, [], [f2615, f770, f702])).
fof(f2615, plain, (~ (e3 = op(e3, e3)) | ~ spl78_20), inference(forward_demodulation, [], [f282, f772])).
fof(f772, plain, ((e3 = op(e2, e3)) | ~ spl78_20), inference(avatar_component_clause, [], [f770])).
fof(f2614, plain, (~ spl78_3 | ~ spl78_51), inference(avatar_split_clause, [], [f2613, f902, f698])).
fof(f2613, plain, (~ (e2 = op(e3, e3)) | ~ spl78_51), inference(forward_demodulation, [], [f279, f904])).
fof(f2559, plain, (~ spl78_57 | ~ spl78_61), inference(avatar_split_clause, [], [f2551, f945, f928])).
fof(f2551, plain, (~ (e0 = op(e0, e1)) | ~ spl78_61), inference(backward_demodulation, [], [f283, f947])).
fof(f2548, plain, (spl78_9 | ~ spl78_2 | ~ spl78_149), inference(avatar_split_clause, [], [f2547, f1670, f694, f724])).
fof(f2547, plain, ((e0 = op(e3, e1)) | (~ spl78_2 | ~ spl78_149)), inference(forward_demodulation, [], [f1671, f696])).
fof(f2539, plain, (~ spl78_21 | ~ spl78_30 | spl78_155), inference(avatar_contradiction_clause, [], [f2538])).
fof(f2538, plain, ($false | (~ spl78_21 | ~ spl78_30 | spl78_155)), inference(subsumption_resolution, [], [f2537, f815])).
fof(f2537, plain, (~ (e1 = op(e2, e0)) | (~ spl78_21 | spl78_155)), inference(forward_demodulation, [], [f1700, f777])).
fof(f2499, plain, (~ spl78_7 | ~ spl78_8), inference(avatar_contradiction_clause, [], [f2498])).
fof(f2498, plain, ($false | (~ spl78_7 | ~ spl78_8)), inference(subsumption_resolution, [], [f2497, f312])).
fof(f312, plain, ~ (e2 = e3), inference(cnf_transformation, [], [f4])).
fof(f2497, plain, ((e2 = e3) | (~ spl78_7 | ~ spl78_8)), inference(backward_demodulation, [], [f721, f717])).
fof(f2489, plain, (~ spl78_13 | ~ spl78_16), inference(avatar_contradiction_clause, [], [f2488])).
fof(f2488, plain, ($false | (~ spl78_13 | ~ spl78_16)), inference(subsumption_resolution, [], [f2486, f309])).
fof(f2486, plain, ((e0 = e3) | (~ spl78_13 | ~ spl78_16)), inference(backward_demodulation, [], [f755, f743])).
fof(f2484, plain, (~ spl78_21 | spl78_29 | ~ spl78_146), inference(avatar_contradiction_clause, [], [f2483])).
fof(f2483, plain, ($false | (~ spl78_21 | spl78_29 | ~ spl78_146)), inference(subsumption_resolution, [], [f2479, f810])).
fof(f810, plain, (~ (e0 = op(e2, e0)) | spl78_29), inference(avatar_component_clause, [], [f809])).
fof(f2479, plain, ((e0 = op(e2, e0)) | (~ spl78_21 | ~ spl78_146)), inference(backward_demodulation, [], [f1657, f777])).
fof(f2477, plain, (~ spl78_27 | ~ spl78_28), inference(avatar_contradiction_clause, [], [f2476])).
fof(f2476, plain, ($false | (~ spl78_27 | ~ spl78_28)), inference(subsumption_resolution, [], [f2475, f312])).
fof(f2475, plain, ((e2 = e3) | (~ spl78_27 | ~ spl78_28)), inference(forward_demodulation, [], [f806, f802])).
fof(f2467, plain, (~ spl78_41 | ~ spl78_44), inference(avatar_contradiction_clause, [], [f2466])).
fof(f2466, plain, ($false | (~ spl78_41 | ~ spl78_44)), inference(subsumption_resolution, [], [f2465, f309])).
fof(f2465, plain, ((e0 = e3) | (~ spl78_41 | ~ spl78_44)), inference(forward_demodulation, [], [f874, f862])).
fof(f2445, plain, (~ spl78_53 | ~ spl78_57), inference(avatar_split_clause, [], [f2438, f928, f911])).
fof(f2438, plain, (~ (e0 = op(e0, e2)) | ~ spl78_57), inference(backward_demodulation, [], [f286, f930])).
fof(f2436, plain, (~ spl78_60 | ~ spl78_62 | spl78_164), inference(avatar_split_clause, [], [f2432, f1744, f949, f940])).
fof(f2432, plain, (~ (e3 = op(e0, e1)) | (~ spl78_62 | spl78_164)), inference(backward_demodulation, [], [f1746, f951])).
fof(f2433, plain, (~ spl78_58 | ~ spl78_62), inference(avatar_split_clause, [], [f2425, f949, f932])).
fof(f2425, plain, (~ (e1 = op(e0, e1)) | ~ spl78_62), inference(backward_demodulation, [], [f283, f951])).
fof(f2423, plain, (~ spl78_11 | ~ spl78_2 | spl78_163), inference(avatar_split_clause, [], [f2422, f1738, f694, f732])).
fof(f2422, plain, (~ (e2 = op(e3, e1)) | (~ spl78_2 | spl78_163)), inference(forward_demodulation, [], [f1740, f696])).
fof(f2416, plain, (~ spl78_57 | ~ spl78_41), inference(avatar_split_clause, [], [f2415, f860, f928])).
fof(f2415, plain, (~ (e0 = op(e0, e1)) | ~ spl78_41), inference(forward_demodulation, [], [f265, f862])).
fof(f2411, plain, (~ spl78_54 | ~ spl78_146 | spl78_148), inference(avatar_split_clause, [], [f2231, f1665, f1656, f915])).
fof(f2231, plain, (~ (e1 = op(e0, e2)) | (~ spl78_146 | spl78_148)), inference(forward_demodulation, [], [f1667, f1657])).
fof(f2406, plain, (~ spl78_48 | ~ spl78_41 | spl78_165), inference(avatar_split_clause, [], [f2330, f1750, f860, f889])).
fof(f2330, plain, (~ (e3 = op(e1, e0)) | (~ spl78_41 | spl78_165)), inference(backward_demodulation, [], [f1752, f862])).
fof(f1752, plain, (~ (e3 = op(e1, op(e1, e1))) | spl78_165), inference(avatar_component_clause, [], [f1750])).
fof(f2382, plain, (~ spl78_10 | ~ spl78_2), inference(avatar_split_clause, [], [f2381, f694, f728])).
fof(f2381, plain, (~ (e1 = op(e3, e1)) | ~ spl78_2), inference(forward_demodulation, [], [f305, f696])).
fof(f2375, plain, (~ spl78_9 | ~ spl78_2 | spl78_149), inference(avatar_split_clause, [], [f2364, f1670, f694, f724])).
fof(f2364, plain, (~ (e0 = op(e3, e1)) | (~ spl78_2 | spl78_149)), inference(backward_demodulation, [], [f1672, f696])).
fof(f2361, plain, (~ spl78_9 | ~ spl78_11), inference(avatar_contradiction_clause, [], [f2360])).
fof(f2360, plain, ($false | (~ spl78_9 | ~ spl78_11)), inference(subsumption_resolution, [], [f2359, f308])).
fof(f2359, plain, ((e0 = e2) | (~ spl78_9 | ~ spl78_11)), inference(forward_demodulation, [], [f734, f726])).
fof(f2349, plain, (~ spl78_29 | ~ spl78_30), inference(avatar_contradiction_clause, [], [f2348])).
fof(f2348, plain, ($false | (~ spl78_29 | ~ spl78_30)), inference(subsumption_resolution, [], [f2347, f307])).
fof(f2347, plain, ((e0 = e1) | (~ spl78_29 | ~ spl78_30)), inference(forward_demodulation, [], [f815, f811])).
fof(f2346, plain, (~ spl78_34 | ~ spl78_35), inference(avatar_contradiction_clause, [], [f2345])).
fof(f2345, plain, ($false | (~ spl78_34 | ~ spl78_35)), inference(subsumption_resolution, [], [f2344, f310])).
fof(f2344, plain, ((e1 = e2) | (~ spl78_34 | ~ spl78_35)), inference(backward_demodulation, [], [f836, f832])).
fof(f2321, plain, (~ spl78_58 | ~ spl78_60), inference(avatar_contradiction_clause, [], [f2320])).
fof(f2320, plain, ($false | (~ spl78_58 | ~ spl78_60)), inference(subsumption_resolution, [], [f2319, f311])).
fof(f2319, plain, ((e1 = e3) | (~ spl78_58 | ~ spl78_60)), inference(backward_demodulation, [], [f942, f934])).
fof(f2311, plain, (~ spl78_17 | spl78_158 | ~ spl78_163), inference(avatar_split_clause, [], [f2309, f1738, f1713, f758])).
fof(f1713, plain, (spl78_158 <=> (e0 = op(op(e3, op(e3, e3)), e3))), introduced(avatar_definition, [new_symbols(naming, [spl78_158])])).
fof(f2309, plain, (~ (e0 = op(e2, e3)) | (spl78_158 | ~ spl78_163)), inference(backward_demodulation, [], [f1715, f1739])).
fof(f1715, plain, (~ (e0 = op(op(e3, op(e3, e3)), e3)) | spl78_158), inference(avatar_component_clause, [], [f1713])).
fof(f2290, plain, (~ spl78_35 | ~ spl78_39), inference(avatar_split_clause, [], [f2289, f851, f834])).
fof(f2289, plain, (~ (e2 = op(e1, e3)) | ~ spl78_39), inference(forward_demodulation, [], [f294, f853])).
fof(f2273, plain, (~ spl78_3 | ~ spl78_6 | ~ spl78_33 | spl78_158), inference(avatar_contradiction_clause, [], [f2272])).
fof(f2272, plain, ($false | (~ spl78_3 | ~ spl78_6 | ~ spl78_33 | spl78_158)), inference(subsumption_resolution, [], [f2270, f828])).
fof(f2270, plain, (~ (e0 = op(e1, e3)) | (~ spl78_3 | ~ spl78_6 | spl78_158)), inference(backward_demodulation, [], [f2161, f713])).
fof(f2161, plain, (~ (e0 = op(op(e3, e2), e3)) | (~ spl78_3 | spl78_158)), inference(backward_demodulation, [], [f1715, f700])).
fof(f2266, plain, (~ spl78_9 | ~ spl78_12), inference(avatar_contradiction_clause, [], [f2265])).
fof(f2265, plain, ($false | (~ spl78_9 | ~ spl78_12)), inference(subsumption_resolution, [], [f2263, f309])).
fof(f2263, plain, ((e0 = e3) | (~ spl78_9 | ~ spl78_12)), inference(backward_demodulation, [], [f738, f726])).
fof(f2249, plain, (~ spl78_22 | ~ spl78_26), inference(avatar_split_clause, [], [f2248, f796, f779])).
fof(f2248, plain, (~ (e1 = op(e2, e2)) | ~ spl78_26), inference(backward_demodulation, [], [f298, f798])).
fof(f2242, plain, (~ spl78_50 | ~ spl78_52), inference(avatar_contradiction_clause, [], [f2241])).
fof(f2241, plain, ($false | (~ spl78_50 | ~ spl78_52)), inference(subsumption_resolution, [], [f2240, f311])).
fof(f2240, plain, ((e1 = e3) | (~ spl78_50 | ~ spl78_52)), inference(forward_demodulation, [], [f908, f900])).
fof(f2239, plain, (~ spl78_62 | ~ spl78_63), inference(avatar_contradiction_clause, [], [f2238])).
fof(f2238, plain, ($false | (~ spl78_62 | ~ spl78_63)), inference(subsumption_resolution, [], [f2237, f310])).
fof(f2237, plain, ((e1 = e2) | (~ spl78_62 | ~ spl78_63)), inference(backward_demodulation, [], [f955, f951])).
fof(f2195, plain, (~ spl78_25 | ~ spl78_29), inference(avatar_split_clause, [], [f2194, f809, f792])).
fof(f2194, plain, (~ (e0 = op(e2, e1)) | ~ spl78_29), inference(forward_demodulation, [], [f295, f811])).
fof(f2185, plain, (~ spl78_23 | ~ spl78_19), inference(avatar_split_clause, [], [f2184, f766, f783])).
fof(f2184, plain, (~ (e2 = op(e2, e2)) | ~ spl78_19), inference(forward_demodulation, [], [f300, f768])).
fof(f2182, plain, (~ spl78_23 | ~ spl78_39), inference(avatar_split_clause, [], [f2181, f851, f783])).
fof(f2181, plain, (~ (e2 = op(e2, e2)) | ~ spl78_39), inference(forward_demodulation, [], [f274, f853])).
fof(f2151, plain, (~ spl78_13 | ~ spl78_14), inference(avatar_contradiction_clause, [], [f2150])).
fof(f2150, plain, ($false | (~ spl78_13 | ~ spl78_14)), inference(subsumption_resolution, [], [f2149, f307])).
fof(f2149, plain, ((e0 = e1) | (~ spl78_13 | ~ spl78_14)), inference(forward_demodulation, [], [f747, f743])).
fof(f2148, plain, (~ spl78_19 | ~ spl78_20), inference(avatar_contradiction_clause, [], [f2147])).
fof(f2147, plain, ($false | (~ spl78_19 | ~ spl78_20)), inference(subsumption_resolution, [], [f2146, f312])).
fof(f2146, plain, ((e2 = e3) | (~ spl78_19 | ~ spl78_20)), inference(backward_demodulation, [], [f772, f768])).
fof(f2138, plain, (~ spl78_42 | ~ spl78_44), inference(avatar_contradiction_clause, [], [f2137])).
fof(f2137, plain, ($false | (~ spl78_42 | ~ spl78_44)), inference(subsumption_resolution, [], [f2136, f311])).
fof(f2136, plain, ((e1 = e3) | (~ spl78_42 | ~ spl78_44)), inference(forward_demodulation, [], [f874, f866])).
fof(f2130, plain, (~ spl78_47 | ~ spl78_50 | ~ spl78_64 | spl78_154), inference(avatar_split_clause, [], [f2126, f1693, f957, f898, f885])).
fof(f2126, plain, (~ (e2 = op(e1, e0)) | (~ spl78_50 | ~ spl78_64 | spl78_154)), inference(backward_demodulation, [], [f2110, f900])).
fof(f2110, plain, (~ (e2 = op(op(e0, e3), e0)) | (~ spl78_64 | spl78_154)), inference(forward_demodulation, [], [f1695, f959])).
fof(f2116, plain, (~ spl78_63 | ~ spl78_64), inference(avatar_contradiction_clause, [], [f2115])).
fof(f2115, plain, ($false | (~ spl78_63 | ~ spl78_64)), inference(subsumption_resolution, [], [f2114, f312])).
fof(f2114, plain, ((e2 = e3) | (~ spl78_63 | ~ spl78_64)), inference(backward_demodulation, [], [f959, f955])).
fof(f2105, plain, (~ spl78_57 | ~ spl78_25), inference(avatar_split_clause, [], [f2104, f792, f928])).
fof(f2104, plain, (~ (e0 = op(e0, e1)) | ~ spl78_25), inference(forward_demodulation, [], [f266, f794])).
fof(f2053, plain, (~ spl78_22 | ~ spl78_25 | ~ spl78_53 | spl78_156), inference(avatar_contradiction_clause, [], [f2052])).
fof(f2052, plain, ($false | (~ spl78_22 | ~ spl78_25 | ~ spl78_53 | spl78_156)), inference(subsumption_resolution, [], [f2049, f913])).
fof(f2049, plain, (~ (e0 = op(e0, e2)) | (~ spl78_22 | ~ spl78_25 | spl78_156)), inference(backward_demodulation, [], [f1912, f794])).
fof(f1912, plain, (~ (e0 = op(op(e2, e1), e2)) | (~ spl78_22 | spl78_156)), inference(forward_demodulation, [], [f1705, f781])).
fof(f2044, plain, (~ spl78_42 | ~ spl78_43), inference(avatar_contradiction_clause, [], [f2043])).
fof(f2043, plain, ($false | (~ spl78_42 | ~ spl78_43)), inference(subsumption_resolution, [], [f2042, f310])).
fof(f2042, plain, ((e1 = e2) | (~ spl78_42 | ~ spl78_43)), inference(backward_demodulation, [], [f870, f866])).
fof(f2033, plain, (~ spl78_62 | ~ spl78_64), inference(avatar_contradiction_clause, [], [f2032])).
fof(f2032, plain, ($false | (~ spl78_62 | ~ spl78_64)), inference(subsumption_resolution, [], [f2031, f311])).
fof(f2031, plain, ((e1 = e3) | (~ spl78_62 | ~ spl78_64)), inference(backward_demodulation, [], [f959, f951])).
fof(f2026, plain, (~ spl78_51 | ~ spl78_64 | spl78_159), inference(avatar_contradiction_clause, [], [f2025])).
fof(f2025, plain, ($false | (~ spl78_51 | ~ spl78_64 | spl78_159)), inference(subsumption_resolution, [], [f2024, f904])).
fof(f2024, plain, (~ (e2 = op(e0, e3)) | (~ spl78_64 | spl78_159)), inference(forward_demodulation, [], [f1720, f959])).
fof(f2009, plain, (~ spl78_17 | ~ spl78_20), inference(avatar_contradiction_clause, [], [f2008])).
fof(f2008, plain, ($false | (~ spl78_17 | ~ spl78_20)), inference(subsumption_resolution, [], [f2007, f309])).
fof(f2007, plain, ((e0 = e3) | (~ spl78_17 | ~ spl78_20)), inference(forward_demodulation, [], [f772, f760])).
fof(f2003, plain, (~ spl78_1 | ~ spl78_33), inference(avatar_split_clause, [], [f2002, f826, f690])).
fof(f2002, plain, (~ (e0 = op(e3, e3)) | ~ spl78_33), inference(forward_demodulation, [], [f281, f828])).
fof(f2000, plain, (~ spl78_9 | ~ spl78_10), inference(avatar_contradiction_clause, [], [f1999])).
fof(f1999, plain, ($false | (~ spl78_9 | ~ spl78_10)), inference(subsumption_resolution, [], [f1998, f307])).
fof(f1998, plain, ((e0 = e1) | (~ spl78_9 | ~ spl78_10)), inference(backward_demodulation, [], [f730, f726])).
fof(f1997, plain, (~ spl78_10 | ~ spl78_12), inference(avatar_contradiction_clause, [], [f1996])).
fof(f1996, plain, ($false | (~ spl78_10 | ~ spl78_12)), inference(subsumption_resolution, [], [f1995, f311])).
fof(f1995, plain, ((e1 = e3) | (~ spl78_10 | ~ spl78_12)), inference(backward_demodulation, [], [f738, f730])).
fof(f1987, plain, (~ spl78_17 | ~ spl78_18), inference(avatar_contradiction_clause, [], [f1986])).
fof(f1986, plain, ($false | (~ spl78_17 | ~ spl78_18)), inference(subsumption_resolution, [], [f1985, f307])).
fof(f1985, plain, ((e0 = e1) | (~ spl78_17 | ~ spl78_18)), inference(backward_demodulation, [], [f764, f760])).
fof(f1983, plain, (~ spl78_18 | ~ spl78_20), inference(avatar_contradiction_clause, [], [f1982])).
fof(f1982, plain, ($false | (~ spl78_18 | ~ spl78_20)), inference(subsumption_resolution, [], [f1981, f311])).
fof(f1981, plain, ((e1 = e3) | (~ spl78_18 | ~ spl78_20)), inference(forward_demodulation, [], [f772, f764])).
fof(f1975, plain, (~ spl78_33 | ~ spl78_34), inference(avatar_contradiction_clause, [], [f1974])).
fof(f1974, plain, ($false | (~ spl78_33 | ~ spl78_34)), inference(subsumption_resolution, [], [f1973, f307])).
fof(f1973, plain, ((e0 = e1) | (~ spl78_33 | ~ spl78_34)), inference(backward_demodulation, [], [f832, f828])).
fof(f1970, plain, (~ spl78_43 | ~ spl78_44), inference(avatar_contradiction_clause, [], [f1969])).
fof(f1969, plain, ($false | (~ spl78_43 | ~ spl78_44)), inference(subsumption_resolution, [], [f1968, f312])).
fof(f1968, plain, ((e2 = e3) | (~ spl78_43 | ~ spl78_44)), inference(forward_demodulation, [], [f874, f870])).
fof(f1967, plain, (~ spl78_14 | ~ spl78_46), inference(avatar_split_clause, [], [f1966, f881, f745])).
fof(f1966, plain, (~ (e1 = op(e3, e0)) | ~ spl78_46), inference(backward_demodulation, [], [f263, f883])).
fof(f1963, plain, (~ spl78_53 | ~ spl78_54), inference(avatar_contradiction_clause, [], [f1962])).
fof(f1962, plain, ($false | (~ spl78_53 | ~ spl78_54)), inference(subsumption_resolution, [], [f1961, f307])).
fof(f1961, plain, ((e0 = e1) | (~ spl78_53 | ~ spl78_54)), inference(forward_demodulation, [], [f917, f913])).
fof(f1957, plain, (~ spl78_37 | ~ spl78_43 | spl78_143), inference(avatar_split_clause, [], [f1956, f1642, f868, f843])).
fof(f1956, plain, (~ (e0 = op(e1, e2)) | (~ spl78_43 | spl78_143)), inference(forward_demodulation, [], [f1644, f870])).
fof(f1953, plain, (~ spl78_50 | ~ spl78_64 | spl78_152), inference(avatar_split_clause, [], [f1952, f1684, f957, f898])).
fof(f1952, plain, (~ (e1 = op(e0, e3)) | (~ spl78_64 | spl78_152)), inference(forward_demodulation, [], [f1686, f959])).
fof(f1949, plain, (~ spl78_9 | ~ spl78_40 | ~ spl78_43 | spl78_162), inference(avatar_split_clause, [], [f1948, f1733, f868, f855, f724])).
fof(f1948, plain, (~ (e0 = op(e3, e1)) | (~ spl78_40 | ~ spl78_43 | spl78_162)), inference(forward_demodulation, [], [f1947, f857])).
fof(f1947, plain, (~ (e0 = op(op(e1, e2), e1)) | (~ spl78_43 | spl78_162)), inference(forward_demodulation, [], [f1735, f870])).
fof(f1917, plain, (~ spl78_18 | ~ spl78_22), inference(avatar_contradiction_clause, [], [f1916])).
fof(f1916, plain, ($false | (~ spl78_18 | ~ spl78_22)), inference(subsumption_resolution, [], [f1915, f781])).
fof(f1915, plain, (~ (e1 = op(e2, e2)) | ~ spl78_18), inference(forward_demodulation, [], [f300, f764])).
fof(f1914, plain, (~ spl78_25 | ~ spl78_22 | spl78_146), inference(avatar_split_clause, [], [f1913, f1656, f779, f792])).
fof(f1913, plain, (~ (e0 = op(e2, e1)) | (~ spl78_22 | spl78_146)), inference(forward_demodulation, [], [f1658, f781])).
fof(f1902, plain, (~ spl78_15 | ~ spl78_1 | spl78_163), inference(avatar_split_clause, [], [f1901, f1738, f690, f749])).
fof(f1901, plain, (~ (e2 = op(e3, e0)) | (~ spl78_1 | spl78_163)), inference(forward_demodulation, [], [f1740, f692])).
fof(f1900, plain, (~ spl78_1 | ~ spl78_2), inference(avatar_contradiction_clause, [], [f1899])).
fof(f1899, plain, ($false | (~ spl78_1 | ~ spl78_2)), inference(subsumption_resolution, [], [f1898, f307])).
fof(f1898, plain, ((e0 = e1) | (~ spl78_1 | ~ spl78_2)), inference(backward_demodulation, [], [f696, f692])).
fof(f1887, plain, (~ spl78_4 | ~ spl78_12), inference(avatar_split_clause, [], [f1885, f736, f702])).
fof(f1885, plain, (~ (e3 = op(e3, e3)) | ~ spl78_12), inference(backward_demodulation, [], [f305, f738])).
fof(f1875, plain, (~ spl78_37 | ~ spl78_18 | ~ spl78_24 | spl78_156), inference(avatar_split_clause, [], [f1873, f1703, f787, f762, f843])).
fof(f1873, plain, (~ (e0 = op(e1, e2)) | (~ spl78_18 | ~ spl78_24 | spl78_156)), inference(backward_demodulation, [], [f1865, f764])).
fof(f1865, plain, (~ (e0 = op(op(e2, e3), e2)) | (~ spl78_24 | spl78_156)), inference(backward_demodulation, [], [f1705, f789])).
fof(f1872, plain, (~ spl78_22 | ~ spl78_24), inference(avatar_contradiction_clause, [], [f1871])).
fof(f1871, plain, ($false | (~ spl78_22 | ~ spl78_24)), inference(subsumption_resolution, [], [f1870, f311])).
fof(f1870, plain, ((e1 = e3) | (~ spl78_22 | ~ spl78_24)), inference(backward_demodulation, [], [f789, f781])).
fof(f1868, plain, (~ spl78_17 | ~ spl78_24 | spl78_146), inference(avatar_split_clause, [], [f1864, f1656, f787, f758])).
fof(f1864, plain, (~ (e0 = op(e2, e3)) | (~ spl78_24 | spl78_146)), inference(backward_demodulation, [], [f1658, f789])).
fof(f1867, plain, (~ spl78_20 | ~ spl78_24), inference(avatar_split_clause, [], [f1863, f787, f770])).
fof(f1863, plain, (~ (e3 = op(e2, e3)) | ~ spl78_24), inference(backward_demodulation, [], [f300, f789])).
fof(f1855, plain, (~ spl78_29 | ~ spl78_31), inference(avatar_contradiction_clause, [], [f1854])).
fof(f1854, plain, ($false | (~ spl78_29 | ~ spl78_31)), inference(subsumption_resolution, [], [f1853, f308])).
fof(f1853, plain, ((e0 = e2) | (~ spl78_29 | ~ spl78_31)), inference(backward_demodulation, [], [f819, f811])).
fof(f1850, plain, (~ spl78_27 | ~ spl78_31), inference(avatar_split_clause, [], [f1846, f817, f800])).
fof(f1846, plain, (~ (e2 = op(e2, e1)) | ~ spl78_31), inference(backward_demodulation, [], [f295, f819])).
fof(f1844, plain, (~ spl78_34 | ~ spl78_36), inference(avatar_contradiction_clause, [], [f1843])).
fof(f1843, plain, ($false | (~ spl78_34 | ~ spl78_36)), inference(subsumption_resolution, [], [f1842, f311])).
fof(f1842, plain, ((e1 = e3) | (~ spl78_34 | ~ spl78_36)), inference(backward_demodulation, [], [f840, f832])).
fof(f1841, plain, (~ spl78_4 | ~ spl78_36), inference(avatar_split_clause, [], [f1839, f838, f702])).
fof(f1839, plain, (~ (e3 = op(e3, e3)) | ~ spl78_36), inference(backward_demodulation, [], [f281, f840])).
fof(f1835, plain, (~ spl78_24 | ~ spl78_40), inference(avatar_split_clause, [], [f1831, f855, f787])).
fof(f1831, plain, (~ (e3 = op(e2, e2)) | ~ spl78_40), inference(backward_demodulation, [], [f274, f857])).
fof(f1829, plain, (~ spl78_41 | ~ spl78_43), inference(avatar_contradiction_clause, [], [f1828])).
fof(f1828, plain, ($false | (~ spl78_41 | ~ spl78_43)), inference(subsumption_resolution, [], [f1827, f308])).
fof(f1827, plain, ((e0 = e2) | (~ spl78_41 | ~ spl78_43)), inference(backward_demodulation, [], [f870, f862])).
fof(f1825, plain, (~ spl78_40 | ~ spl78_43 | spl78_165), inference(avatar_split_clause, [], [f1820, f1750, f868, f855])).
fof(f1820, plain, (~ (e3 = op(e1, e2)) | (~ spl78_43 | spl78_165)), inference(backward_demodulation, [], [f1752, f870])).
fof(f1824, plain, (~ spl78_35 | ~ spl78_43), inference(avatar_split_clause, [], [f1817, f868, f834])).
fof(f1817, plain, (~ (e2 = op(e1, e3)) | ~ spl78_43), inference(backward_demodulation, [], [f293, f870])).
fof(f1822, plain, (~ spl78_11 | ~ spl78_43), inference(avatar_split_clause, [], [f1815, f868, f732])).
fof(f1815, plain, (~ (e2 = op(e3, e1)) | ~ spl78_43), inference(backward_demodulation, [], [f269, f870])).
fof(f1813, plain, (~ spl78_33 | ~ spl78_45), inference(avatar_split_clause, [], [f1808, f877, f826])).
fof(f1808, plain, (~ (e0 = op(e1, e3)) | ~ spl78_45), inference(backward_demodulation, [], [f291, f879])).
fof(f1810, plain, (~ spl78_13 | ~ spl78_45), inference(avatar_split_clause, [], [f1805, f877, f741])).
fof(f1805, plain, (~ (e0 = op(e3, e0)) | ~ spl78_45), inference(backward_demodulation, [], [f263, f879])).
fof(f1809, plain, (~ spl78_29 | ~ spl78_45), inference(avatar_split_clause, [], [f1804, f877, f809])).
fof(f1804, plain, (~ (e0 = op(e2, e0)) | ~ spl78_45), inference(backward_demodulation, [], [f262, f879])).
fof(f1802, plain, (~ spl78_30 | ~ spl78_51 | ~ spl78_64 | spl78_160), inference(avatar_split_clause, [], [f1797, f1723, f957, f902, f813])).
fof(f1797, plain, (~ (e1 = op(e2, e0)) | (~ spl78_51 | ~ spl78_64 | spl78_160)), inference(backward_demodulation, [], [f1768, f904])).
fof(f1768, plain, (~ (e1 = op(op(e0, e3), e0)) | (~ spl78_64 | spl78_160)), inference(backward_demodulation, [], [f1725, f959])).
fof(f1792, plain, (~ spl78_49 | ~ spl78_53), inference(avatar_split_clause, [], [f1788, f911, f894])).
fof(f1788, plain, (~ (e0 = op(e0, e3)) | ~ spl78_53), inference(backward_demodulation, [], [f288, f913])).
fof(f1789, plain, (~ spl78_37 | ~ spl78_53), inference(avatar_split_clause, [], [f1785, f911, f843])).
fof(f1785, plain, (~ (e0 = op(e1, e2)) | ~ spl78_53), inference(backward_demodulation, [], [f271, f913])).
fof(f1773, plain, (~ spl78_56 | ~ spl78_64), inference(avatar_split_clause, [], [f1765, f957, f923])).
fof(f1765, plain, (~ (e3 = op(e0, e2)) | ~ spl78_64), inference(backward_demodulation, [], [f284, f959])).
fof(f1772, plain, (~ spl78_60 | ~ spl78_64), inference(avatar_split_clause, [], [f1764, f957, f940])).
fof(f1764, plain, (~ (e3 = op(e0, e1)) | ~ spl78_64), inference(backward_demodulation, [], [f283, f959])).
fof(f1770, plain, (~ spl78_32 | ~ spl78_64), inference(avatar_split_clause, [], [f1762, f957, f821])).
fof(f1762, plain, (~ (e3 = op(e2, e0)) | ~ spl78_64), inference(backward_demodulation, [], [f260, f959])).
fof(f1769, plain, (~ spl78_48 | ~ spl78_64), inference(avatar_split_clause, [], [f1761, f957, f889])).
fof(f1761, plain, (~ (e3 = op(e1, e0)) | ~ spl78_64), inference(backward_demodulation, [], [f259, f959])).
fof(f1760, plain, (~ spl78_166 | ~ spl78_22 | ~ spl78_156), inference(avatar_split_clause, [], [f688, f1703, f779, f1756])).
fof(f688, plain, (~ (e0 = op(op(e2, op(e2, e2)), e2)) | ~ (e1 = op(e2, e2)) | ~ (e3 = op(e2, op(e2, e2)))), inference(cnf_transformation, [], [f53])).
fof(f53, plain, (~ (e0 = op(op(e2, op(e2, e2)), e2)) | ~ (e1 = op(e2, e2)) | ~ (e3 = op(e2, op(e2, e2)))), inference(ennf_transformation, [], [f29])).
fof(f29, plain, ~ ((e0 = op(op(e2, op(e2, e2)), e2)) & (e1 = op(e2, e2)) & (e3 = op(e2, op(e2, e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG160+1.p', ax29)).
fof(f1759, plain, (~ spl78_166 | ~ spl78_21 | ~ spl78_148), inference(avatar_split_clause, [], [f687, f1665, f775, f1756])).
fof(f687, plain, (~ (e1 = op(op(e2, op(e2, e2)), e2)) | ~ (e0 = op(e2, e2)) | ~ (e3 = op(e2, op(e2, e2)))), inference(cnf_transformation, [], [f52])).
fof(f52, plain, (~ (e1 = op(op(e2, op(e2, e2)), e2)) | ~ (e0 = op(e2, e2)) | ~ (e3 = op(e2, op(e2, e2)))), inference(ennf_transformation, [], [f28])).
fof(f28, plain, ~ ((e1 = op(op(e2, op(e2, e2)), e2)) & (e0 = op(e2, e2)) & (e3 = op(e2, op(e2, e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG160+1.p', ax28)).
fof(f1754, plain, (~ spl78_165 | ~ spl78_43 | ~ spl78_162), inference(avatar_split_clause, [], [f686, f1733, f868, f1750])).
fof(f686, plain, (~ (e0 = op(op(e1, op(e1, e1)), e1)) | ~ (e2 = op(e1, e1)) | ~ (e3 = op(e1, op(e1, e1)))), inference(cnf_transformation, [], [f51])).
fof(f51, plain, (~ (e0 = op(op(e1, op(e1, e1)), e1)) | ~ (e2 = op(e1, e1)) | ~ (e3 = op(e1, op(e1, e1)))), inference(ennf_transformation, [], [f27])).
fof(f27, plain, ~ ((e0 = op(op(e1, op(e1, e1)), e1)) & (e2 = op(e1, e1)) & (e3 = op(e1, op(e1, e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG160+1.p', ax27)).
fof(f1753, plain, (~ spl78_165 | ~ spl78_41 | ~ spl78_145), inference(avatar_split_clause, [], [f685, f1651, f860, f1750])).
fof(f685, plain, (~ (e2 = op(op(e1, op(e1, e1)), e1)) | ~ (e0 = op(e1, e1)) | ~ (e3 = op(e1, op(e1, e1)))), inference(cnf_transformation, [], [f50])).
fof(f50, plain, (~ (e2 = op(op(e1, op(e1, e1)), e1)) | ~ (e0 = op(e1, e1)) | ~ (e3 = op(e1, op(e1, e1)))), inference(ennf_transformation, [], [f26])).
fof(f26, plain, ~ ((e2 = op(op(e1, op(e1, e1)), e1)) & (e0 = op(e1, e1)) & (e3 = op(e1, op(e1, e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG160+1.p', ax26)).
fof(f1748, plain, (~ spl78_164 | ~ spl78_63 | ~ spl78_160), inference(avatar_split_clause, [], [f684, f1723, f953, f1744])).
fof(f684, plain, (~ (e1 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e2) | ~ (e3 = op(e0, op(e0, e0)))), inference(cnf_transformation, [], [f49])).
fof(f49, plain, (~ (e1 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e2) | ~ (e3 = op(e0, op(e0, e0)))), inference(ennf_transformation, [], [f25])).
fof(f25, plain, ~ ((e1 = op(op(e0, op(e0, e0)), e0)) & (op(e0, e0) = e2) & (e3 = op(e0, op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG160+1.p', ax25)).
fof(f1747, plain, (~ spl78_164 | ~ spl78_62 | ~ spl78_154), inference(avatar_split_clause, [], [f683, f1693, f949, f1744])).
fof(f683, plain, (~ (e2 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e1) | ~ (e3 = op(e0, op(e0, e0)))), inference(cnf_transformation, [], [f48])).
fof(f48, plain, (~ (e2 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e1) | ~ (e3 = op(e0, op(e0, e0)))), inference(ennf_transformation, [], [f24])).
fof(f24, plain, ~ ((e2 = op(op(e0, op(e0, e0)), e0)) & (op(e0, e0) = e1) & (e3 = op(e0, op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG160+1.p', ax24)).
fof(f1742, plain, (~ spl78_163 | ~ spl78_2 | ~ spl78_158), inference(avatar_split_clause, [], [f682, f1713, f694, f1738])).
fof(f682, plain, (~ (e0 = op(op(e3, op(e3, e3)), e3)) | ~ (e1 = op(e3, e3)) | ~ (e2 = op(e3, op(e3, e3)))), inference(cnf_transformation, [], [f47])).
fof(f47, plain, (~ (e0 = op(op(e3, op(e3, e3)), e3)) | ~ (e1 = op(e3, e3)) | ~ (e2 = op(e3, op(e3, e3)))), inference(ennf_transformation, [], [f23])).
fof(f23, plain, ~ ((e0 = op(op(e3, op(e3, e3)), e3)) & (e1 = op(e3, e3)) & (e2 = op(e3, op(e3, e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG160+1.p', ax23)).
fof(f1741, plain, (~ spl78_163 | ~ spl78_1 | ~ spl78_151), inference(avatar_split_clause, [], [f681, f1679, f690, f1738])).
fof(f681, plain, (~ (e1 = op(op(e3, op(e3, e3)), e3)) | ~ (e0 = op(e3, e3)) | ~ (e2 = op(e3, op(e3, e3)))), inference(cnf_transformation, [], [f46])).
fof(f46, plain, (~ (e1 = op(op(e3, op(e3, e3)), e3)) | ~ (e0 = op(e3, e3)) | ~ (e2 = op(e3, op(e3, e3)))), inference(ennf_transformation, [], [f22])).
fof(f22, plain, ~ ((e1 = op(op(e3, op(e3, e3)), e3)) & (e0 = op(e3, e3)) & (e2 = op(e3, op(e3, e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG160+1.p', ax22)).
fof(f1736, plain, (~ spl78_161 | ~ spl78_44 | ~ spl78_162), inference(avatar_split_clause, [], [f680, f1733, f872, f1728])).
fof(f680, plain, (~ (e0 = op(op(e1, op(e1, e1)), e1)) | ~ (e3 = op(e1, e1)) | ~ (e2 = op(e1, op(e1, e1)))), inference(cnf_transformation, [], [f45])).
fof(f45, plain, (~ (e0 = op(op(e1, op(e1, e1)), e1)) | ~ (e3 = op(e1, e1)) | ~ (e2 = op(e1, op(e1, e1)))), inference(ennf_transformation, [], [f21])).
fof(f21, plain, ~ ((e0 = op(op(e1, op(e1, e1)), e1)) & (e3 = op(e1, e1)) & (e2 = op(e1, op(e1, e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG160+1.p', ax21)).
fof(f1731, plain, (~ spl78_161 | ~ spl78_41 | ~ spl78_144), inference(avatar_split_clause, [], [f679, f1646, f860, f1728])).
fof(f679, plain, (~ (e3 = op(op(e1, op(e1, e1)), e1)) | ~ (e0 = op(e1, e1)) | ~ (e2 = op(e1, op(e1, e1)))), inference(cnf_transformation, [], [f44])).
fof(f44, plain, (~ (e3 = op(op(e1, op(e1, e1)), e1)) | ~ (e0 = op(e1, e1)) | ~ (e2 = op(e1, op(e1, e1)))), inference(ennf_transformation, [], [f20])).
fof(f20, plain, ~ ((e3 = op(op(e1, op(e1, e1)), e1)) & (e0 = op(e1, e1)) & (e2 = op(e1, op(e1, e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG160+1.p', ax20)).
fof(f1726, plain, (~ spl78_159 | ~ spl78_64 | ~ spl78_160), inference(avatar_split_clause, [], [f678, f1723, f957, f1718])).
fof(f678, plain, (~ (e1 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e3) | ~ (e2 = op(e0, op(e0, e0)))), inference(cnf_transformation, [], [f43])).
fof(f43, plain, (~ (e1 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e3) | ~ (e2 = op(e0, op(e0, e0)))), inference(ennf_transformation, [], [f19])).
fof(f19, plain, ~ ((e1 = op(op(e0, op(e0, e0)), e0)) & (op(e0, e0) = e3) & (e2 = op(e0, op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG160+1.p', ax19)).
fof(f1721, plain, (~ spl78_159 | ~ spl78_62 | ~ spl78_153), inference(avatar_split_clause, [], [f677, f1688, f949, f1718])).
fof(f677, plain, (~ (e3 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e1) | ~ (e2 = op(e0, op(e0, e0)))), inference(cnf_transformation, [], [f42])).
fof(f42, plain, (~ (e3 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e1) | ~ (e2 = op(e0, op(e0, e0)))), inference(ennf_transformation, [], [f18])).
fof(f18, plain, ~ ((e3 = op(op(e0, op(e0, e0)), e0)) & (op(e0, e0) = e1) & (e2 = op(e0, op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG160+1.p', ax18)).
fof(f1716, plain, (~ spl78_157 | ~ spl78_3 | ~ spl78_158), inference(avatar_split_clause, [], [f676, f1713, f698, f1708])).
fof(f676, plain, (~ (e0 = op(op(e3, op(e3, e3)), e3)) | ~ (e2 = op(e3, e3)) | ~ (e1 = op(e3, op(e3, e3)))), inference(cnf_transformation, [], [f41])).
fof(f41, plain, (~ (e0 = op(op(e3, op(e3, e3)), e3)) | ~ (e2 = op(e3, e3)) | ~ (e1 = op(e3, op(e3, e3)))), inference(ennf_transformation, [], [f17])).
fof(f17, plain, ~ ((e0 = op(op(e3, op(e3, e3)), e3)) & (e2 = op(e3, e3)) & (e1 = op(e3, op(e3, e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG160+1.p', ax17)).
fof(f1711, plain, (~ spl78_157 | ~ spl78_1 | ~ spl78_150), inference(avatar_split_clause, [], [f675, f1674, f690, f1708])).
fof(f675, plain, (~ (e2 = op(op(e3, op(e3, e3)), e3)) | ~ (e0 = op(e3, e3)) | ~ (e1 = op(e3, op(e3, e3)))), inference(cnf_transformation, [], [f40])).
fof(f40, plain, (~ (e2 = op(op(e3, op(e3, e3)), e3)) | ~ (e0 = op(e3, e3)) | ~ (e1 = op(e3, op(e3, e3)))), inference(ennf_transformation, [], [f16])).
fof(f16, plain, ~ ((e2 = op(op(e3, op(e3, e3)), e3)) & (e0 = op(e3, e3)) & (e1 = op(e3, op(e3, e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG160+1.p', ax16)).
fof(f1706, plain, (~ spl78_155 | ~ spl78_24 | ~ spl78_156), inference(avatar_split_clause, [], [f674, f1703, f787, f1698])).
fof(f674, plain, (~ (e0 = op(op(e2, op(e2, e2)), e2)) | ~ (e3 = op(e2, e2)) | ~ (e1 = op(e2, op(e2, e2)))), inference(cnf_transformation, [], [f39])).
fof(f39, plain, (~ (e0 = op(op(e2, op(e2, e2)), e2)) | ~ (e3 = op(e2, e2)) | ~ (e1 = op(e2, op(e2, e2)))), inference(ennf_transformation, [], [f15])).
fof(f15, plain, ~ ((e0 = op(op(e2, op(e2, e2)), e2)) & (e3 = op(e2, e2)) & (e1 = op(e2, op(e2, e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG160+1.p', ax15)).
fof(f1701, plain, (~ spl78_155 | ~ spl78_21 | ~ spl78_147), inference(avatar_split_clause, [], [f673, f1660, f775, f1698])).
fof(f673, plain, (~ (e3 = op(op(e2, op(e2, e2)), e2)) | ~ (e0 = op(e2, e2)) | ~ (e1 = op(e2, op(e2, e2)))), inference(cnf_transformation, [], [f38])).
fof(f38, plain, (~ (e3 = op(op(e2, op(e2, e2)), e2)) | ~ (e0 = op(e2, e2)) | ~ (e1 = op(e2, op(e2, e2)))), inference(ennf_transformation, [], [f14])).
fof(f14, plain, ~ ((e3 = op(op(e2, op(e2, e2)), e2)) & (e0 = op(e2, e2)) & (e1 = op(e2, op(e2, e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG160+1.p', ax14)).
fof(f1696, plain, (~ spl78_152 | ~ spl78_64 | ~ spl78_154), inference(avatar_split_clause, [], [f672, f1693, f957, f1684])).
fof(f672, plain, (~ (e2 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e3) | ~ (e1 = op(e0, op(e0, e0)))), inference(cnf_transformation, [], [f37])).
fof(f37, plain, (~ (e2 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e3) | ~ (e1 = op(e0, op(e0, e0)))), inference(ennf_transformation, [], [f13])).
fof(f13, plain, ~ ((e2 = op(op(e0, op(e0, e0)), e0)) & (op(e0, e0) = e3) & (e1 = op(e0, op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG160+1.p', ax13)).
fof(f1691, plain, (~ spl78_152 | ~ spl78_63 | ~ spl78_153), inference(avatar_split_clause, [], [f671, f1688, f953, f1684])).
fof(f671, plain, (~ (e3 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e2) | ~ (e1 = op(e0, op(e0, e0)))), inference(cnf_transformation, [], [f36])).
fof(f36, plain, (~ (e3 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e2) | ~ (e1 = op(e0, op(e0, e0)))), inference(ennf_transformation, [], [f12])).
fof(f12, plain, ~ ((e3 = op(op(e0, op(e0, e0)), e0)) & (op(e0, e0) = e2) & (e1 = op(e0, op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG160+1.p', ax12)).
fof(f1682, plain, (~ spl78_149 | ~ spl78_3 | ~ spl78_151), inference(avatar_split_clause, [], [f670, f1679, f698, f1670])).
fof(f670, plain, (~ (e1 = op(op(e3, op(e3, e3)), e3)) | ~ (e2 = op(e3, e3)) | ~ (e0 = op(e3, op(e3, e3)))), inference(cnf_transformation, [], [f35])).
fof(f35, plain, (~ (e1 = op(op(e3, op(e3, e3)), e3)) | ~ (e2 = op(e3, e3)) | ~ (e0 = op(e3, op(e3, e3)))), inference(ennf_transformation, [], [f11])).
fof(f11, plain, ~ ((e1 = op(op(e3, op(e3, e3)), e3)) & (e2 = op(e3, e3)) & (e0 = op(e3, op(e3, e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG160+1.p', ax11)).
fof(f1677, plain, (~ spl78_149 | ~ spl78_2 | ~ spl78_150), inference(avatar_split_clause, [], [f669, f1674, f694, f1670])).
fof(f669, plain, (~ (e2 = op(op(e3, op(e3, e3)), e3)) | ~ (e1 = op(e3, e3)) | ~ (e0 = op(e3, op(e3, e3)))), inference(cnf_transformation, [], [f34])).
fof(f34, plain, (~ (e2 = op(op(e3, op(e3, e3)), e3)) | ~ (e1 = op(e3, e3)) | ~ (e0 = op(e3, op(e3, e3)))), inference(ennf_transformation, [], [f10])).
fof(f10, plain, ~ ((e2 = op(op(e3, op(e3, e3)), e3)) & (e1 = op(e3, e3)) & (e0 = op(e3, op(e3, e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG160+1.p', ax10)).
fof(f1668, plain, (~ spl78_146 | ~ spl78_24 | ~ spl78_148), inference(avatar_split_clause, [], [f668, f1665, f787, f1656])).
fof(f668, plain, (~ (e1 = op(op(e2, op(e2, e2)), e2)) | ~ (e3 = op(e2, e2)) | ~ (e0 = op(e2, op(e2, e2)))), inference(cnf_transformation, [], [f33])).
fof(f33, plain, (~ (e1 = op(op(e2, op(e2, e2)), e2)) | ~ (e3 = op(e2, e2)) | ~ (e0 = op(e2, op(e2, e2)))), inference(ennf_transformation, [], [f9])).
fof(f9, plain, ~ ((e1 = op(op(e2, op(e2, e2)), e2)) & (e3 = op(e2, e2)) & (e0 = op(e2, op(e2, e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG160+1.p', ax9)).
fof(f1663, plain, (~ spl78_146 | ~ spl78_22 | ~ spl78_147), inference(avatar_split_clause, [], [f667, f1660, f779, f1656])).
fof(f667, plain, (~ (e3 = op(op(e2, op(e2, e2)), e2)) | ~ (e1 = op(e2, e2)) | ~ (e0 = op(e2, op(e2, e2)))), inference(cnf_transformation, [], [f32])).
fof(f32, plain, (~ (e3 = op(op(e2, op(e2, e2)), e2)) | ~ (e1 = op(e2, e2)) | ~ (e0 = op(e2, op(e2, e2)))), inference(ennf_transformation, [], [f8])).
fof(f8, plain, ~ ((e3 = op(op(e2, op(e2, e2)), e2)) & (e1 = op(e2, e2)) & (e0 = op(e2, op(e2, e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG160+1.p', ax8)).
fof(f1654, plain, (~ spl78_143 | ~ spl78_44 | ~ spl78_145), inference(avatar_split_clause, [], [f666, f1651, f872, f1642])).
fof(f666, plain, (~ (e2 = op(op(e1, op(e1, e1)), e1)) | ~ (e3 = op(e1, e1)) | ~ (e0 = op(e1, op(e1, e1)))), inference(cnf_transformation, [], [f31])).
fof(f31, plain, (~ (e2 = op(op(e1, op(e1, e1)), e1)) | ~ (e3 = op(e1, e1)) | ~ (e0 = op(e1, op(e1, e1)))), inference(ennf_transformation, [], [f7])).
fof(f7, plain, ~ ((e2 = op(op(e1, op(e1, e1)), e1)) & (e3 = op(e1, e1)) & (e0 = op(e1, op(e1, e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG160+1.p', ax7)).
fof(f1649, plain, (~ spl78_143 | ~ spl78_43 | ~ spl78_144), inference(avatar_split_clause, [], [f665, f1646, f868, f1642])).
fof(f665, plain, (~ (e3 = op(op(e1, op(e1, e1)), e1)) | ~ (e2 = op(e1, e1)) | ~ (e0 = op(e1, op(e1, e1)))), inference(cnf_transformation, [], [f30])).
fof(f30, plain, (~ (e3 = op(op(e1, op(e1, e1)), e1)) | ~ (e2 = op(e1, e1)) | ~ (e0 = op(e1, op(e1, e1)))), inference(ennf_transformation, [], [f6])).
fof(f6, plain, ~ ((e3 = op(op(e1, op(e1, e1)), e1)) & (e2 = op(e1, e1)) & (e0 = op(e1, op(e1, e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG160+1.p', ax6)).
fof(f1640, plain, (spl78_142 | spl78_141 | spl78_140 | spl78_139 | spl78_138 | spl78_137 | spl78_136 | spl78_135 | spl78_134 | spl78_133 | spl78_132 | spl78_131 | spl78_130 | spl78_129 | spl78_128 | spl78_127 | spl78_126 | spl78_125 | spl78_124 | spl78_123 | spl78_122 | spl78_121 | spl78_120 | spl78_119 | spl78_118 | spl78_117 | spl78_116 | spl78_115 | spl78_114 | spl78_113 | spl78_112 | spl78_111 | spl78_110 | spl78_109 | spl78_108 | spl78_107 | spl78_106 | spl78_105 | spl78_104 | spl78_103 | spl78_102 | spl78_101 | spl78_100 | spl78_99 | spl78_98 | spl78_97 | spl78_96 | spl78_95 | spl78_94 | spl78_93 | spl78_92 | spl78_91 | spl78_90 | spl78_89 | spl78_88 | spl78_87 | spl78_86 | spl78_85 | spl78_84 | spl78_83 | spl78_82 | spl78_81 | spl78_80 | spl78_4), inference(avatar_split_clause, [], [f655, f702, f1129, f1137, f1145, f1153, f1161, f1169, f1177, f1185, f1193, f1201, f1209, f1217, f1225, f1233, f1241, f1249, f1257, f1265, f1273, f1281, f1289, f1297, f1305, f1313, f1321, f1329, f1337, f1345, f1353, f1361, f1369, f1377, f1385, f1393, f1401, f1409, f1417, f1425, f1433, f1441, f1449, f1457, f1465, f1473, f1481, f1489, f1497, f1505, f1513, f1521, f1529, f1537, f1545, f1553, f1561, f1569, f1577, f1585, f1593, f1601, f1609, f1617, f1625])).
fof(f1625, plain, (spl78_142 <=> sP0), introduced(avatar_definition, [new_symbols(naming, [spl78_142])])).
fof(f1617, plain, (spl78_141 <=> sP1), introduced(avatar_definition, [new_symbols(naming, [spl78_141])])).
fof(f1609, plain, (spl78_140 <=> sP2), introduced(avatar_definition, [new_symbols(naming, [spl78_140])])).
fof(f1601, plain, (spl78_139 <=> sP3), introduced(avatar_definition, [new_symbols(naming, [spl78_139])])).
fof(f1593, plain, (spl78_138 <=> sP4), introduced(avatar_definition, [new_symbols(naming, [spl78_138])])).
fof(f1585, plain, (spl78_137 <=> sP5), introduced(avatar_definition, [new_symbols(naming, [spl78_137])])).
fof(f1577, plain, (spl78_136 <=> sP6), introduced(avatar_definition, [new_symbols(naming, [spl78_136])])).
fof(f1569, plain, (spl78_135 <=> sP7), introduced(avatar_definition, [new_symbols(naming, [spl78_135])])).
fof(f1561, plain, (spl78_134 <=> sP8), introduced(avatar_definition, [new_symbols(naming, [spl78_134])])).
fof(f1553, plain, (spl78_133 <=> sP9), introduced(avatar_definition, [new_symbols(naming, [spl78_133])])).
fof(f1545, plain, (spl78_132 <=> sP10), introduced(avatar_definition, [new_symbols(naming, [spl78_132])])).
fof(f1537, plain, (spl78_131 <=> sP11), introduced(avatar_definition, [new_symbols(naming, [spl78_131])])).
fof(f1529, plain, (spl78_130 <=> sP12), introduced(avatar_definition, [new_symbols(naming, [spl78_130])])).
fof(f1521, plain, (spl78_129 <=> sP13), introduced(avatar_definition, [new_symbols(naming, [spl78_129])])).
fof(f1513, plain, (spl78_128 <=> sP14), introduced(avatar_definition, [new_symbols(naming, [spl78_128])])).
fof(f1505, plain, (spl78_127 <=> sP15), introduced(avatar_definition, [new_symbols(naming, [spl78_127])])).
fof(f1497, plain, (spl78_126 <=> sP16), introduced(avatar_definition, [new_symbols(naming, [spl78_126])])).
fof(f1489, plain, (spl78_125 <=> sP17), introduced(avatar_definition, [new_symbols(naming, [spl78_125])])).
fof(f1481, plain, (spl78_124 <=> sP18), introduced(avatar_definition, [new_symbols(naming, [spl78_124])])).
fof(f1473, plain, (spl78_123 <=> sP19), introduced(avatar_definition, [new_symbols(naming, [spl78_123])])).
fof(f1465, plain, (spl78_122 <=> sP20), introduced(avatar_definition, [new_symbols(naming, [spl78_122])])).
fof(f1457, plain, (spl78_121 <=> sP21), introduced(avatar_definition, [new_symbols(naming, [spl78_121])])).
fof(f1449, plain, (spl78_120 <=> sP22), introduced(avatar_definition, [new_symbols(naming, [spl78_120])])).
fof(f1441, plain, (spl78_119 <=> sP23), introduced(avatar_definition, [new_symbols(naming, [spl78_119])])).
fof(f1433, plain, (spl78_118 <=> sP24), introduced(avatar_definition, [new_symbols(naming, [spl78_118])])).
fof(f1425, plain, (spl78_117 <=> sP25), introduced(avatar_definition, [new_symbols(naming, [spl78_117])])).
fof(f1417, plain, (spl78_116 <=> sP26), introduced(avatar_definition, [new_symbols(naming, [spl78_116])])).
fof(f1409, plain, (spl78_115 <=> sP27), introduced(avatar_definition, [new_symbols(naming, [spl78_115])])).
fof(f1401, plain, (spl78_114 <=> sP28), introduced(avatar_definition, [new_symbols(naming, [spl78_114])])).
fof(f1393, plain, (spl78_113 <=> sP29), introduced(avatar_definition, [new_symbols(naming, [spl78_113])])).
fof(f1385, plain, (spl78_112 <=> sP30), introduced(avatar_definition, [new_symbols(naming, [spl78_112])])).
fof(f1377, plain, (spl78_111 <=> sP31), introduced(avatar_definition, [new_symbols(naming, [spl78_111])])).
fof(f1369, plain, (spl78_110 <=> sP32), introduced(avatar_definition, [new_symbols(naming, [spl78_110])])).
fof(f1361, plain, (spl78_109 <=> sP33), introduced(avatar_definition, [new_symbols(naming, [spl78_109])])).
fof(f1353, plain, (spl78_108 <=> sP34), introduced(avatar_definition, [new_symbols(naming, [spl78_108])])).
fof(f1345, plain, (spl78_107 <=> sP35), introduced(avatar_definition, [new_symbols(naming, [spl78_107])])).
fof(f1337, plain, (spl78_106 <=> sP36), introduced(avatar_definition, [new_symbols(naming, [spl78_106])])).
fof(f1329, plain, (spl78_105 <=> sP37), introduced(avatar_definition, [new_symbols(naming, [spl78_105])])).
fof(f1321, plain, (spl78_104 <=> sP38), introduced(avatar_definition, [new_symbols(naming, [spl78_104])])).
fof(f1313, plain, (spl78_103 <=> sP39), introduced(avatar_definition, [new_symbols(naming, [spl78_103])])).
fof(f1305, plain, (spl78_102 <=> sP40), introduced(avatar_definition, [new_symbols(naming, [spl78_102])])).
fof(f1297, plain, (spl78_101 <=> sP41), introduced(avatar_definition, [new_symbols(naming, [spl78_101])])).
fof(f1289, plain, (spl78_100 <=> sP42), introduced(avatar_definition, [new_symbols(naming, [spl78_100])])).
fof(f1281, plain, (spl78_99 <=> sP43), introduced(avatar_definition, [new_symbols(naming, [spl78_99])])).
fof(f1273, plain, (spl78_98 <=> sP44), introduced(avatar_definition, [new_symbols(naming, [spl78_98])])).
fof(f1265, plain, (spl78_97 <=> sP45), introduced(avatar_definition, [new_symbols(naming, [spl78_97])])).
fof(f1257, plain, (spl78_96 <=> sP46), introduced(avatar_definition, [new_symbols(naming, [spl78_96])])).
fof(f1249, plain, (spl78_95 <=> sP47), introduced(avatar_definition, [new_symbols(naming, [spl78_95])])).
fof(f1241, plain, (spl78_94 <=> sP48), introduced(avatar_definition, [new_symbols(naming, [spl78_94])])).
fof(f1233, plain, (spl78_93 <=> sP49), introduced(avatar_definition, [new_symbols(naming, [spl78_93])])).
fof(f1225, plain, (spl78_92 <=> sP50), introduced(avatar_definition, [new_symbols(naming, [spl78_92])])).
fof(f1217, plain, (spl78_91 <=> sP51), introduced(avatar_definition, [new_symbols(naming, [spl78_91])])).
fof(f1209, plain, (spl78_90 <=> sP52), introduced(avatar_definition, [new_symbols(naming, [spl78_90])])).
fof(f1201, plain, (spl78_89 <=> sP53), introduced(avatar_definition, [new_symbols(naming, [spl78_89])])).
fof(f1193, plain, (spl78_88 <=> sP54), introduced(avatar_definition, [new_symbols(naming, [spl78_88])])).
fof(f1185, plain, (spl78_87 <=> sP55), introduced(avatar_definition, [new_symbols(naming, [spl78_87])])).
fof(f1177, plain, (spl78_86 <=> sP56), introduced(avatar_definition, [new_symbols(naming, [spl78_86])])).
fof(f1169, plain, (spl78_85 <=> sP57), introduced(avatar_definition, [new_symbols(naming, [spl78_85])])).
fof(f1161, plain, (spl78_84 <=> sP58), introduced(avatar_definition, [new_symbols(naming, [spl78_84])])).
fof(f1153, plain, (spl78_83 <=> sP59), introduced(avatar_definition, [new_symbols(naming, [spl78_83])])).
fof(f1145, plain, (spl78_82 <=> sP60), introduced(avatar_definition, [new_symbols(naming, [spl78_82])])).
fof(f1137, plain, (spl78_81 <=> sP61), introduced(avatar_definition, [new_symbols(naming, [spl78_81])])).
fof(f1129, plain, (spl78_80 <=> sP62), introduced(avatar_definition, [new_symbols(naming, [spl78_80])])).
fof(f655, plain, ((e3 = op(e3, e3)) | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f132])).
fof(f132, plain, (((((e3 = op(e3, e3)) | ~ (e3 = op(e3, e3))) & ((e3 = op(e2, e2)) | ~ (e2 = op(e2, e3))) & ((e3 = op(e1, e1)) | ~ (e1 = op(e1, e3))) & ((op(e0, e0) = e3) | ~ (e0 = op(e0, e3))) & (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63) & ((~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0)), inference(definition_folding, [], [f5, e131, e130, e129, e128, e127, e126, e125, e124, e123, e122, e121, e120, e119, e118, e117, e116, e115, e114, e113, e112, e111, e110, e109, e108, e107, e106, e105, e104, e103, e102, e101, e100, e99, e98, e97, e96, e95, e94, e93, e92, e91, e90, e89, e88, e87, e86, e85, e84, e83, e82, e81, e80, e79, e78, e77, e76, e75, e74, e73, e72, e71, e70, e69, e68, e67, e66, e65, e64, e63, e62, e61, e60, e59, e58, e57, e56, e55, e54])).
fof(f54, plain, ((~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP0), inference(usedef, [], [e54])).
fof(e54, plain, (sP0 <=> (~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f55, plain, ((~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP1), inference(usedef, [], [e55])).
fof(e55, plain, (sP1 <=> (~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f56, plain, ((~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP2), inference(usedef, [], [e56])).
fof(e56, plain, (sP2 <=> (~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f57, plain, ((~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP3), inference(usedef, [], [e57])).
fof(e57, plain, (sP3 <=> (~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f58, plain, ((~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | ~ sP4), inference(usedef, [], [e58])).
fof(e58, plain, (sP4 <=> (~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f59, plain, ((~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | ~ sP5), inference(usedef, [], [e59])).
fof(e59, plain, (sP5 <=> (~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f60, plain, ((~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | ~ sP6), inference(usedef, [], [e60])).
fof(e60, plain, (sP6 <=> (~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f61, plain, ((~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | ~ sP7), inference(usedef, [], [e61])).
fof(e61, plain, (sP7 <=> (~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f62, plain, ((~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | ~ sP8), inference(usedef, [], [e62])).
fof(e62, plain, (sP8 <=> (~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f63, plain, ((~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | ~ sP9), inference(usedef, [], [e63])).
fof(e63, plain, (sP9 <=> (~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f64, plain, ((~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | ~ sP10), inference(usedef, [], [e64])).
fof(e64, plain, (sP10 <=> (~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f65, plain, ((~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | ~ sP11), inference(usedef, [], [e65])).
fof(e65, plain, (sP11 <=> (~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f66, plain, ((~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | ~ sP12), inference(usedef, [], [e66])).
fof(e66, plain, (sP12 <=> (~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f67, plain, ((~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | ~ sP13), inference(usedef, [], [e67])).
fof(e67, plain, (sP13 <=> (~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f68, plain, ((~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | ~ sP14), inference(usedef, [], [e68])).
fof(e68, plain, (sP14 <=> (~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f69, plain, ((~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | ~ sP15), inference(usedef, [], [e69])).
fof(e69, plain, (sP15 <=> (~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f70, plain, ((~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | ~ sP16), inference(usedef, [], [e70])).
fof(e70, plain, (sP16 <=> (~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f71, plain, ((~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | ~ sP17), inference(usedef, [], [e71])).
fof(e71, plain, (sP17 <=> (~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f72, plain, ((~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | ~ sP18), inference(usedef, [], [e72])).
fof(e72, plain, (sP18 <=> (~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f73, plain, ((~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | ~ sP19), inference(usedef, [], [e73])).
fof(e73, plain, (sP19 <=> (~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP19])])).
fof(f74, plain, ((~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP20), inference(usedef, [], [e74])).
fof(e74, plain, (sP20 <=> (~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP20])])).
fof(f75, plain, ((~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP21), inference(usedef, [], [e75])).
fof(e75, plain, (sP21 <=> (~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP21])])).
fof(f76, plain, ((~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP22), inference(usedef, [], [e76])).
fof(e76, plain, (sP22 <=> (~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP22])])).
fof(f77, plain, ((~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP23), inference(usedef, [], [e77])).
fof(e77, plain, (sP23 <=> (~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP23])])).
fof(f78, plain, ((~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | ~ sP24), inference(usedef, [], [e78])).
fof(e78, plain, (sP24 <=> (~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP24])])).
fof(f79, plain, ((~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | ~ sP25), inference(usedef, [], [e79])).
fof(e79, plain, (sP25 <=> (~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP25])])).
fof(f80, plain, ((~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | ~ sP26), inference(usedef, [], [e80])).
fof(e80, plain, (sP26 <=> (~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP26])])).
fof(f81, plain, ((~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | ~ sP27), inference(usedef, [], [e81])).
fof(e81, plain, (sP27 <=> (~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP27])])).
fof(f82, plain, ((~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | ~ sP28), inference(usedef, [], [e82])).
fof(e82, plain, (sP28 <=> (~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP28])])).
fof(f83, plain, ((~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | ~ sP29), inference(usedef, [], [e83])).
fof(e83, plain, (sP29 <=> (~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP29])])).
fof(f84, plain, ((~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | ~ sP30), inference(usedef, [], [e84])).
fof(e84, plain, (sP30 <=> (~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP30])])).
fof(f85, plain, ((~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | ~ sP31), inference(usedef, [], [e85])).
fof(e85, plain, (sP31 <=> (~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP31])])).
fof(f86, plain, ((~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | ~ sP32), inference(usedef, [], [e86])).
fof(e86, plain, (sP32 <=> (~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP32])])).
fof(f87, plain, ((~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | ~ sP33), inference(usedef, [], [e87])).
fof(e87, plain, (sP33 <=> (~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP33])])).
fof(f88, plain, ((~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | ~ sP34), inference(usedef, [], [e88])).
fof(e88, plain, (sP34 <=> (~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP34])])).
fof(f89, plain, ((~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | ~ sP35), inference(usedef, [], [e89])).
fof(e89, plain, (sP35 <=> (~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP35])])).
fof(f90, plain, ((~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | ~ sP36), inference(usedef, [], [e90])).
fof(e90, plain, (sP36 <=> (~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP36])])).
fof(f91, plain, ((~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | ~ sP37), inference(usedef, [], [e91])).
fof(e91, plain, (sP37 <=> (~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP37])])).
fof(f92, plain, ((~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | ~ sP38), inference(usedef, [], [e92])).
fof(e92, plain, (sP38 <=> (~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP38])])).
fof(f93, plain, ((~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | ~ sP39), inference(usedef, [], [e93])).
fof(e93, plain, (sP39 <=> (~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP39])])).
fof(f94, plain, ((~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP40), inference(usedef, [], [e94])).
fof(e94, plain, (sP40 <=> (~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP40])])).
fof(f95, plain, ((~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP41), inference(usedef, [], [e95])).
fof(e95, plain, (sP41 <=> (~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP41])])).
fof(f96, plain, ((~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP42), inference(usedef, [], [e96])).
fof(e96, plain, (sP42 <=> (~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP42])])).
fof(f97, plain, ((~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP43), inference(usedef, [], [e97])).
fof(e97, plain, (sP43 <=> (~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP43])])).
fof(f98, plain, ((~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | ~ sP44), inference(usedef, [], [e98])).
fof(e98, plain, (sP44 <=> (~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP44])])).
fof(f99, plain, ((~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | ~ sP45), inference(usedef, [], [e99])).
fof(e99, plain, (sP45 <=> (~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP45])])).
fof(f100, plain, ((~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | ~ sP46), inference(usedef, [], [e100])).
fof(e100, plain, (sP46 <=> (~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP46])])).
fof(f101, plain, ((~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | ~ sP47), inference(usedef, [], [e101])).
fof(e101, plain, (sP47 <=> (~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP47])])).
fof(f102, plain, ((~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | ~ sP48), inference(usedef, [], [e102])).
fof(e102, plain, (sP48 <=> (~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP48])])).
fof(f103, plain, ((~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | ~ sP49), inference(usedef, [], [e103])).
fof(e103, plain, (sP49 <=> (~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP49])])).
fof(f104, plain, ((~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | ~ sP50), inference(usedef, [], [e104])).
fof(e104, plain, (sP50 <=> (~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP50])])).
fof(f105, plain, ((~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | ~ sP51), inference(usedef, [], [e105])).
fof(e105, plain, (sP51 <=> (~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP51])])).
fof(f106, plain, ((~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | ~ sP52), inference(usedef, [], [e106])).
fof(e106, plain, (sP52 <=> (~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP52])])).
fof(f107, plain, ((~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | ~ sP53), inference(usedef, [], [e107])).
fof(e107, plain, (sP53 <=> (~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP53])])).
fof(f108, plain, ((~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | ~ sP54), inference(usedef, [], [e108])).
fof(e108, plain, (sP54 <=> (~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP54])])).
fof(f109, plain, ((~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | ~ sP55), inference(usedef, [], [e109])).
fof(e109, plain, (sP55 <=> (~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP55])])).
fof(f110, plain, ((~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | ~ sP56), inference(usedef, [], [e110])).
fof(e110, plain, (sP56 <=> (~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP56])])).
fof(f111, plain, ((~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | ~ sP57), inference(usedef, [], [e111])).
fof(e111, plain, (sP57 <=> (~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP57])])).
fof(f112, plain, ((~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | ~ sP58), inference(usedef, [], [e112])).
fof(e112, plain, (sP58 <=> (~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP58])])).
fof(f113, plain, ((~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | ~ sP59), inference(usedef, [], [e113])).
fof(e113, plain, (sP59 <=> (~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP59])])).
fof(f114, plain, ((~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ~ sP60), inference(usedef, [], [e114])).
fof(e114, plain, (sP60 <=> (~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP60])])).
fof(f115, plain, ((~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ~ sP61), inference(usedef, [], [e115])).
fof(e115, plain, (sP61 <=> (~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP61])])).
fof(f116, plain, ((~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ~ sP62), inference(usedef, [], [e116])).
fof(e116, plain, (sP62 <=> (~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP62])])).
fof(f117, plain, ((((e0 = op(e3, e3)) | ~ (e3 = op(e3, e0))) & ((e0 = op(e2, e2)) | ~ (e2 = op(e2, e0))) & ((e0 = op(e1, e1)) | ~ (e1 = op(e1, e0))) & ((e0 = op(e0, e0)) | ~ (e0 = op(e0, e0))) & (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP63), inference(usedef, [], [e117])).
fof(e117, plain, (sP63 <=> (((e0 = op(e3, e3)) | ~ (e3 = op(e3, e0))) & ((e0 = op(e2, e2)) | ~ (e2 = op(e2, e0))) & ((e0 = op(e1, e1)) | ~ (e1 = op(e1, e0))) & ((e0 = op(e0, e0)) | ~ (e0 = op(e0, e0))) & (e0 = op(e0, e0)) & (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP63])])).
fof(f118, plain, ((((e1 = op(e3, e3)) | ~ (e3 = op(e3, e1))) & ((e1 = op(e2, e2)) | ~ (e2 = op(e2, e1))) & ((e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & ((op(e0, e0) = e1) | ~ (e0 = op(e0, e1))) & (e0 = op(e1, e0)) & (e0 = op(e0, e1))) | ~ sP64), inference(usedef, [], [e118])).
fof(e118, plain, (sP64 <=> (((e1 = op(e3, e3)) | ~ (e3 = op(e3, e1))) & ((e1 = op(e2, e2)) | ~ (e2 = op(e2, e1))) & ((e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & ((op(e0, e0) = e1) | ~ (e0 = op(e0, e1))) & (e0 = op(e1, e0)) & (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP64])])).
fof(f119, plain, ((((e2 = op(e3, e3)) | ~ (e3 = op(e3, e2))) & ((e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & ((e2 = op(e1, e1)) | ~ (e1 = op(e1, e2))) & ((op(e0, e0) = e2) | ~ (e0 = op(e0, e2))) & (e0 = op(e2, e0)) & (e0 = op(e0, e2))) | ~ sP65), inference(usedef, [], [e119])).
fof(e119, plain, (sP65 <=> (((e2 = op(e3, e3)) | ~ (e3 = op(e3, e2))) & ((e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & ((e2 = op(e1, e1)) | ~ (e1 = op(e1, e2))) & ((op(e0, e0) = e2) | ~ (e0 = op(e0, e2))) & (e0 = op(e2, e0)) & (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP65])])).
fof(f120, plain, ((((e3 = op(e3, e3)) | ~ (e3 = op(e3, e3))) & ((e3 = op(e2, e2)) | ~ (e2 = op(e2, e3))) & ((e3 = op(e1, e1)) | ~ (e1 = op(e1, e3))) & ((op(e0, e0) = e3) | ~ (e0 = op(e0, e3))) & (e0 = op(e3, e0)) & (e0 = op(e0, e3))) | ~ sP66), inference(usedef, [], [e120])).
fof(e120, plain, (sP66 <=> (((e3 = op(e3, e3)) | ~ (e3 = op(e3, e3))) & ((e3 = op(e2, e2)) | ~ (e2 = op(e2, e3))) & ((e3 = op(e1, e1)) | ~ (e1 = op(e1, e3))) & ((op(e0, e0) = e3) | ~ (e0 = op(e0, e3))) & (e0 = op(e3, e0)) & (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP66])])).
fof(f121, plain, ((((e0 = op(e3, e3)) | ~ (e3 = op(e3, e0))) & ((e0 = op(e2, e2)) | ~ (e2 = op(e2, e0))) & ((e0 = op(e1, e1)) | ~ (e1 = op(e1, e0))) & ((e0 = op(e0, e0)) | ~ (e0 = op(e0, e0))) & (e1 = op(e0, e1)) & (e1 = op(e1, e0))) | ~ sP67), inference(usedef, [], [e121])).
fof(e121, plain, (sP67 <=> (((e0 = op(e3, e3)) | ~ (e3 = op(e3, e0))) & ((e0 = op(e2, e2)) | ~ (e2 = op(e2, e0))) & ((e0 = op(e1, e1)) | ~ (e1 = op(e1, e0))) & ((e0 = op(e0, e0)) | ~ (e0 = op(e0, e0))) & (e1 = op(e0, e1)) & (e1 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP67])])).
fof(f122, plain, ((((e1 = op(e3, e3)) | ~ (e3 = op(e3, e1))) & ((e1 = op(e2, e2)) | ~ (e2 = op(e2, e1))) & ((e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & ((op(e0, e0) = e1) | ~ (e0 = op(e0, e1))) & (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP68), inference(usedef, [], [e122])).
fof(e122, plain, (sP68 <=> (((e1 = op(e3, e3)) | ~ (e3 = op(e3, e1))) & ((e1 = op(e2, e2)) | ~ (e2 = op(e2, e1))) & ((e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & ((op(e0, e0) = e1) | ~ (e0 = op(e0, e1))) & (e1 = op(e1, e1)) & (e1 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP68])])).
fof(f123, plain, ((((e2 = op(e3, e3)) | ~ (e3 = op(e3, e2))) & ((e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & ((e2 = op(e1, e1)) | ~ (e1 = op(e1, e2))) & ((op(e0, e0) = e2) | ~ (e0 = op(e0, e2))) & (e1 = op(e2, e1)) & (e1 = op(e1, e2))) | ~ sP69), inference(usedef, [], [e123])).
fof(e123, plain, (sP69 <=> (((e2 = op(e3, e3)) | ~ (e3 = op(e3, e2))) & ((e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & ((e2 = op(e1, e1)) | ~ (e1 = op(e1, e2))) & ((op(e0, e0) = e2) | ~ (e0 = op(e0, e2))) & (e1 = op(e2, e1)) & (e1 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP69])])).
fof(f124, plain, ((((e3 = op(e3, e3)) | ~ (e3 = op(e3, e3))) & ((e3 = op(e2, e2)) | ~ (e2 = op(e2, e3))) & ((e3 = op(e1, e1)) | ~ (e1 = op(e1, e3))) & ((op(e0, e0) = e3) | ~ (e0 = op(e0, e3))) & (e1 = op(e3, e1)) & (e1 = op(e1, e3))) | ~ sP70), inference(usedef, [], [e124])).
fof(e124, plain, (sP70 <=> (((e3 = op(e3, e3)) | ~ (e3 = op(e3, e3))) & ((e3 = op(e2, e2)) | ~ (e2 = op(e2, e3))) & ((e3 = op(e1, e1)) | ~ (e1 = op(e1, e3))) & ((op(e0, e0) = e3) | ~ (e0 = op(e0, e3))) & (e1 = op(e3, e1)) & (e1 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP70])])).
fof(f125, plain, ((((e0 = op(e3, e3)) | ~ (e3 = op(e3, e0))) & ((e0 = op(e2, e2)) | ~ (e2 = op(e2, e0))) & ((e0 = op(e1, e1)) | ~ (e1 = op(e1, e0))) & ((e0 = op(e0, e0)) | ~ (e0 = op(e0, e0))) & (e2 = op(e0, e2)) & (e2 = op(e2, e0))) | ~ sP71), inference(usedef, [], [e125])).
fof(e125, plain, (sP71 <=> (((e0 = op(e3, e3)) | ~ (e3 = op(e3, e0))) & ((e0 = op(e2, e2)) | ~ (e2 = op(e2, e0))) & ((e0 = op(e1, e1)) | ~ (e1 = op(e1, e0))) & ((e0 = op(e0, e0)) | ~ (e0 = op(e0, e0))) & (e2 = op(e0, e2)) & (e2 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP71])])).
fof(f126, plain, ((((e1 = op(e3, e3)) | ~ (e3 = op(e3, e1))) & ((e1 = op(e2, e2)) | ~ (e2 = op(e2, e1))) & ((e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & ((op(e0, e0) = e1) | ~ (e0 = op(e0, e1))) & (e2 = op(e1, e2)) & (e2 = op(e2, e1))) | ~ sP72), inference(usedef, [], [e126])).
fof(e126, plain, (sP72 <=> (((e1 = op(e3, e3)) | ~ (e3 = op(e3, e1))) & ((e1 = op(e2, e2)) | ~ (e2 = op(e2, e1))) & ((e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & ((op(e0, e0) = e1) | ~ (e0 = op(e0, e1))) & (e2 = op(e1, e2)) & (e2 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP72])])).
fof(f127, plain, ((((e2 = op(e3, e3)) | ~ (e3 = op(e3, e2))) & ((e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & ((e2 = op(e1, e1)) | ~ (e1 = op(e1, e2))) & ((op(e0, e0) = e2) | ~ (e0 = op(e0, e2))) & (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP73), inference(usedef, [], [e127])).
fof(e127, plain, (sP73 <=> (((e2 = op(e3, e3)) | ~ (e3 = op(e3, e2))) & ((e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & ((e2 = op(e1, e1)) | ~ (e1 = op(e1, e2))) & ((op(e0, e0) = e2) | ~ (e0 = op(e0, e2))) & (e2 = op(e2, e2)) & (e2 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP73])])).
fof(f128, plain, ((((e3 = op(e3, e3)) | ~ (e3 = op(e3, e3))) & ((e3 = op(e2, e2)) | ~ (e2 = op(e2, e3))) & ((e3 = op(e1, e1)) | ~ (e1 = op(e1, e3))) & ((op(e0, e0) = e3) | ~ (e0 = op(e0, e3))) & (e2 = op(e3, e2)) & (e2 = op(e2, e3))) | ~ sP74), inference(usedef, [], [e128])).
fof(e128, plain, (sP74 <=> (((e3 = op(e3, e3)) | ~ (e3 = op(e3, e3))) & ((e3 = op(e2, e2)) | ~ (e2 = op(e2, e3))) & ((e3 = op(e1, e1)) | ~ (e1 = op(e1, e3))) & ((op(e0, e0) = e3) | ~ (e0 = op(e0, e3))) & (e2 = op(e3, e2)) & (e2 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP74])])).
fof(f129, plain, ((((e0 = op(e3, e3)) | ~ (e3 = op(e3, e0))) & ((e0 = op(e2, e2)) | ~ (e2 = op(e2, e0))) & ((e0 = op(e1, e1)) | ~ (e1 = op(e1, e0))) & ((e0 = op(e0, e0)) | ~ (e0 = op(e0, e0))) & (e3 = op(e0, e3)) & (e3 = op(e3, e0))) | ~ sP75), inference(usedef, [], [e129])).
fof(e129, plain, (sP75 <=> (((e0 = op(e3, e3)) | ~ (e3 = op(e3, e0))) & ((e0 = op(e2, e2)) | ~ (e2 = op(e2, e0))) & ((e0 = op(e1, e1)) | ~ (e1 = op(e1, e0))) & ((e0 = op(e0, e0)) | ~ (e0 = op(e0, e0))) & (e3 = op(e0, e3)) & (e3 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP75])])).
fof(f130, plain, ((((e1 = op(e3, e3)) | ~ (e3 = op(e3, e1))) & ((e1 = op(e2, e2)) | ~ (e2 = op(e2, e1))) & ((e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & ((op(e0, e0) = e1) | ~ (e0 = op(e0, e1))) & (e3 = op(e1, e3)) & (e3 = op(e3, e1))) | ~ sP76), inference(usedef, [], [e130])).
fof(e130, plain, (sP76 <=> (((e1 = op(e3, e3)) | ~ (e3 = op(e3, e1))) & ((e1 = op(e2, e2)) | ~ (e2 = op(e2, e1))) & ((e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & ((op(e0, e0) = e1) | ~ (e0 = op(e0, e1))) & (e3 = op(e1, e3)) & (e3 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP76])])).
fof(f131, plain, ((((e2 = op(e3, e3)) | ~ (e3 = op(e3, e2))) & ((e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & ((e2 = op(e1, e1)) | ~ (e1 = op(e1, e2))) & ((op(e0, e0) = e2) | ~ (e0 = op(e0, e2))) & (e3 = op(e2, e3)) & (e3 = op(e3, e2))) | ~ sP77), inference(usedef, [], [e131])).
fof(e131, plain, (sP77 <=> (((e2 = op(e3, e3)) | ~ (e3 = op(e3, e2))) & ((e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & ((e2 = op(e1, e1)) | ~ (e1 = op(e1, e2))) & ((op(e0, e0) = e2) | ~ (e0 = op(e0, e2))) & (e3 = op(e2, e3)) & (e3 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP77])])).
fof(f5, plain, (((((e3 = op(e3, e3)) | ~ (e3 = op(e3, e3))) & ((e3 = op(e2, e2)) | ~ (e2 = op(e2, e3))) & ((e3 = op(e1, e1)) | ~ (e1 = op(e1, e3))) & ((op(e0, e0) = e3) | ~ (e0 = op(e0, e3))) & (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | (((e2 = op(e3, e3)) | ~ (e3 = op(e3, e2))) & ((e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & ((e2 = op(e1, e1)) | ~ (e1 = op(e1, e2))) & ((op(e0, e0) = e2) | ~ (e0 = op(e0, e2))) & (e3 = op(e2, e3)) & (e3 = op(e3, e2))) | (((e1 = op(e3, e3)) | ~ (e3 = op(e3, e1))) & ((e1 = op(e2, e2)) | ~ (e2 = op(e2, e1))) & ((e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & ((op(e0, e0) = e1) | ~ (e0 = op(e0, e1))) & (e3 = op(e1, e3)) & (e3 = op(e3, e1))) | (((e0 = op(e3, e3)) | ~ (e3 = op(e3, e0))) & ((e0 = op(e2, e2)) | ~ (e2 = op(e2, e0))) & ((e0 = op(e1, e1)) | ~ (e1 = op(e1, e0))) & ((e0 = op(e0, e0)) | ~ (e0 = op(e0, e0))) & (e3 = op(e0, e3)) & (e3 = op(e3, e0))) | (((e3 = op(e3, e3)) | ~ (e3 = op(e3, e3))) & ((e3 = op(e2, e2)) | ~ (e2 = op(e2, e3))) & ((e3 = op(e1, e1)) | ~ (e1 = op(e1, e3))) & ((op(e0, e0) = e3) | ~ (e0 = op(e0, e3))) & (e2 = op(e3, e2)) & (e2 = op(e2, e3))) | (((e2 = op(e3, e3)) | ~ (e3 = op(e3, e2))) & ((e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & ((e2 = op(e1, e1)) | ~ (e1 = op(e1, e2))) & ((op(e0, e0) = e2) | ~ (e0 = op(e0, e2))) & (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | (((e1 = op(e3, e3)) | ~ (e3 = op(e3, e1))) & ((e1 = op(e2, e2)) | ~ (e2 = op(e2, e1))) & ((e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & ((op(e0, e0) = e1) | ~ (e0 = op(e0, e1))) & (e2 = op(e1, e2)) & (e2 = op(e2, e1))) | (((e0 = op(e3, e3)) | ~ (e3 = op(e3, e0))) & ((e0 = op(e2, e2)) | ~ (e2 = op(e2, e0))) & ((e0 = op(e1, e1)) | ~ (e1 = op(e1, e0))) & ((e0 = op(e0, e0)) | ~ (e0 = op(e0, e0))) & (e2 = op(e0, e2)) & (e2 = op(e2, e0))) | (((e3 = op(e3, e3)) | ~ (e3 = op(e3, e3))) & ((e3 = op(e2, e2)) | ~ (e2 = op(e2, e3))) & ((e3 = op(e1, e1)) | ~ (e1 = op(e1, e3))) & ((op(e0, e0) = e3) | ~ (e0 = op(e0, e3))) & (e1 = op(e3, e1)) & (e1 = op(e1, e3))) | (((e2 = op(e3, e3)) | ~ (e3 = op(e3, e2))) & ((e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & ((e2 = op(e1, e1)) | ~ (e1 = op(e1, e2))) & ((op(e0, e0) = e2) | ~ (e0 = op(e0, e2))) & (e1 = op(e2, e1)) & (e1 = op(e1, e2))) | (((e1 = op(e3, e3)) | ~ (e3 = op(e3, e1))) & ((e1 = op(e2, e2)) | ~ (e2 = op(e2, e1))) & ((e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & ((op(e0, e0) = e1) | ~ (e0 = op(e0, e1))) & (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | (((e0 = op(e3, e3)) | ~ (e3 = op(e3, e0))) & ((e0 = op(e2, e2)) | ~ (e2 = op(e2, e0))) & ((e0 = op(e1, e1)) | ~ (e1 = op(e1, e0))) & ((e0 = op(e0, e0)) | ~ (e0 = op(e0, e0))) & (e1 = op(e0, e1)) & (e1 = op(e1, e0))) | (((e3 = op(e3, e3)) | ~ (e3 = op(e3, e3))) & ((e3 = op(e2, e2)) | ~ (e2 = op(e2, e3))) & ((e3 = op(e1, e1)) | ~ (e1 = op(e1, e3))) & ((op(e0, e0) = e3) | ~ (e0 = op(e0, e3))) & (e0 = op(e3, e0)) & (e0 = op(e0, e3))) | (((e2 = op(e3, e3)) | ~ (e3 = op(e3, e2))) & ((e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & ((e2 = op(e1, e1)) | ~ (e1 = op(e1, e2))) & ((op(e0, e0) = e2) | ~ (e0 = op(e0, e2))) & (e0 = op(e2, e0)) & (e0 = op(e0, e2))) | (((e1 = op(e3, e3)) | ~ (e3 = op(e3, e1))) & ((e1 = op(e2, e2)) | ~ (e2 = op(e2, e1))) & ((e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & ((op(e0, e0) = e1) | ~ (e0 = op(e0, e1))) & (e0 = op(e1, e0)) & (e0 = op(e0, e1))) | (((e0 = op(e3, e3)) | ~ (e3 = op(e3, e0))) & ((e0 = op(e2, e2)) | ~ (e2 = op(e2, e0))) & ((e0 = op(e1, e1)) | ~ (e1 = op(e1, e0))) & ((e0 = op(e0, e0)) | ~ (e0 = op(e0, e0))) & (e0 = op(e0, e0)) & (e0 = op(e0, e0)))) & ((~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | (~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | (~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | (~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | (~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | (~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | (~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | (~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | (~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | (~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | (~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | (~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | (~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | (~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | (~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | (~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | (~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | (~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | (~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | (~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | (~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | (~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | (~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | (~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | (~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | (~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | (~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | (~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | (~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | (~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | (~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | (~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | (~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | (~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | (~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | (~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | (~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | (~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | (~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | (~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | (~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | (~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | (~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | (~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | (~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | (~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | (~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | (~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | (~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | (~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | (~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | (~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | (~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | (~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | (~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | (~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | (~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | (~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | (~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | (~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | (~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | (~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | (~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | (~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG160+1.p', ax5)).
fof(f1639, plain, (spl78_142 | spl78_141 | spl78_140 | spl78_139 | spl78_138 | spl78_137 | spl78_136 | spl78_135 | spl78_134 | spl78_133 | spl78_132 | spl78_131 | spl78_130 | spl78_129 | spl78_128 | spl78_127 | spl78_126 | spl78_125 | spl78_124 | spl78_123 | spl78_122 | spl78_121 | spl78_120 | spl78_119 | spl78_118 | spl78_117 | spl78_116 | spl78_115 | spl78_114 | spl78_113 | spl78_112 | spl78_111 | spl78_110 | spl78_109 | spl78_108 | spl78_107 | spl78_106 | spl78_105 | spl78_104 | spl78_103 | spl78_102 | spl78_101 | spl78_100 | spl78_99 | spl78_98 | spl78_97 | spl78_96 | spl78_95 | spl78_94 | spl78_93 | spl78_92 | spl78_91 | spl78_90 | spl78_89 | spl78_88 | spl78_87 | spl78_86 | spl78_85 | spl78_84 | spl78_83 | spl78_82 | spl78_81 | spl78_80 | ~ spl78_4), inference(avatar_split_clause, [], [f656, f702, f1129, f1137, f1145, f1153, f1161, f1169, f1177, f1185, f1193, f1201, f1209, f1217, f1225, f1233, f1241, f1249, f1257, f1265, f1273, f1281, f1289, f1297, f1305, f1313, f1321, f1329, f1337, f1345, f1353, f1361, f1369, f1377, f1385, f1393, f1401, f1409, f1417, f1425, f1433, f1441, f1449, f1457, f1465, f1473, f1481, f1489, f1497, f1505, f1513, f1521, f1529, f1537, f1545, f1553, f1561, f1569, f1577, f1585, f1593, f1601, f1609, f1617, f1625])).
fof(f656, plain, (~ (e3 = op(e3, e3)) | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f132])).
fof(f1636, plain, (spl78_79 | spl78_78 | spl78_77 | spl78_76 | spl78_75 | spl78_74 | spl78_73 | spl78_72 | spl78_71 | spl78_70 | spl78_69 | spl78_68 | spl78_67 | spl78_66 | spl78_65 | spl78_4), inference(avatar_split_clause, [], [f659, f702, f994, f1003, f1012, f1021, f1030, f1039, f1048, f1057, f1066, f1075, f1084, f1093, f1102, f1111, f1120])).
fof(f1120, plain, (spl78_79 <=> sP63), introduced(avatar_definition, [new_symbols(naming, [spl78_79])])).
fof(f1111, plain, (spl78_78 <=> sP64), introduced(avatar_definition, [new_symbols(naming, [spl78_78])])).
fof(f1102, plain, (spl78_77 <=> sP65), introduced(avatar_definition, [new_symbols(naming, [spl78_77])])).
fof(f1093, plain, (spl78_76 <=> sP66), introduced(avatar_definition, [new_symbols(naming, [spl78_76])])).
fof(f1084, plain, (spl78_75 <=> sP67), introduced(avatar_definition, [new_symbols(naming, [spl78_75])])).
fof(f1075, plain, (spl78_74 <=> sP68), introduced(avatar_definition, [new_symbols(naming, [spl78_74])])).
fof(f1066, plain, (spl78_73 <=> sP69), introduced(avatar_definition, [new_symbols(naming, [spl78_73])])).
fof(f1057, plain, (spl78_72 <=> sP70), introduced(avatar_definition, [new_symbols(naming, [spl78_72])])).
fof(f1048, plain, (spl78_71 <=> sP71), introduced(avatar_definition, [new_symbols(naming, [spl78_71])])).
fof(f1039, plain, (spl78_70 <=> sP72), introduced(avatar_definition, [new_symbols(naming, [spl78_70])])).
fof(f1030, plain, (spl78_69 <=> sP73), introduced(avatar_definition, [new_symbols(naming, [spl78_69])])).
fof(f1021, plain, (spl78_68 <=> sP74), introduced(avatar_definition, [new_symbols(naming, [spl78_68])])).
fof(f1012, plain, (spl78_67 <=> sP75), introduced(avatar_definition, [new_symbols(naming, [spl78_67])])).
fof(f1003, plain, (spl78_66 <=> sP76), introduced(avatar_definition, [new_symbols(naming, [spl78_66])])).
fof(f994, plain, (spl78_65 <=> sP77), introduced(avatar_definition, [new_symbols(naming, [spl78_65])])).
fof(f659, plain, ((e3 = op(e3, e3)) | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63), inference(cnf_transformation, [], [f132])).
fof(f1634, plain, (spl78_79 | spl78_78 | spl78_77 | spl78_76 | spl78_75 | spl78_74 | spl78_73 | spl78_72 | spl78_71 | spl78_70 | spl78_69 | spl78_68 | spl78_67 | spl78_66 | spl78_65 | ~ spl78_49 | spl78_64), inference(avatar_split_clause, [], [f661, f957, f894, f994, f1003, f1012, f1021, f1030, f1039, f1048, f1057, f1066, f1075, f1084, f1093, f1102, f1111, f1120])).
fof(f661, plain, ((op(e0, e0) = e3) | ~ (e0 = op(e0, e3)) | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63), inference(cnf_transformation, [], [f132])).
fof(f1633, plain, (spl78_79 | spl78_78 | spl78_77 | spl78_76 | spl78_75 | spl78_74 | spl78_73 | spl78_72 | spl78_71 | spl78_70 | spl78_69 | spl78_68 | spl78_67 | spl78_66 | spl78_65 | ~ spl78_34 | spl78_44), inference(avatar_split_clause, [], [f662, f872, f830, f994, f1003, f1012, f1021, f1030, f1039, f1048, f1057, f1066, f1075, f1084, f1093, f1102, f1111, f1120])).
fof(f662, plain, ((e3 = op(e1, e1)) | ~ (e1 = op(e1, e3)) | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63), inference(cnf_transformation, [], [f132])).
fof(f1632, plain, (spl78_79 | spl78_78 | spl78_77 | spl78_76 | spl78_75 | spl78_74 | spl78_73 | spl78_72 | spl78_71 | spl78_70 | spl78_69 | spl78_68 | spl78_67 | spl78_66 | spl78_65 | ~ spl78_19 | spl78_24), inference(avatar_split_clause, [], [f663, f787, f766, f994, f1003, f1012, f1021, f1030, f1039, f1048, f1057, f1066, f1075, f1084, f1093, f1102, f1111, f1120])).
fof(f663, plain, ((e3 = op(e2, e2)) | ~ (e2 = op(e2, e3)) | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63), inference(cnf_transformation, [], [f132])).
fof(f1631, plain, (~ spl78_142 | spl78_61), inference(avatar_split_clause, [], [f651, f945, f1625])).
fof(f651, plain, ((e0 = op(e0, e0)) | ~ sP0), inference(cnf_transformation, [], [f210])).
fof(f210, plain, ((~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP0), inference(nnf_transformation, [], [f54])).
fof(f1630, plain, (~ spl78_142 | ~ spl78_61), inference(avatar_split_clause, [], [f652, f945, f1625])).
fof(f652, plain, (~ (e0 = op(e0, e0)) | ~ sP0), inference(cnf_transformation, [], [f210])).
fof(f1623, plain, (~ spl78_141 | spl78_61), inference(avatar_split_clause, [], [f647, f945, f1617])).
fof(f647, plain, ((e0 = op(e0, e0)) | ~ sP1), inference(cnf_transformation, [], [f209])).
fof(f209, plain, ((~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP1), inference(nnf_transformation, [], [f55])).
fof(f1622, plain, (~ spl78_141 | ~ spl78_61), inference(avatar_split_clause, [], [f648, f945, f1617])).
fof(f648, plain, (~ (e0 = op(e0, e0)) | ~ sP1), inference(cnf_transformation, [], [f209])).
fof(f1615, plain, (~ spl78_140 | spl78_61), inference(avatar_split_clause, [], [f643, f945, f1609])).
fof(f643, plain, ((e0 = op(e0, e0)) | ~ sP2), inference(cnf_transformation, [], [f208])).
fof(f208, plain, ((~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP2), inference(nnf_transformation, [], [f56])).
fof(f1614, plain, (~ spl78_140 | ~ spl78_61), inference(avatar_split_clause, [], [f644, f945, f1609])).
fof(f644, plain, (~ (e0 = op(e0, e0)) | ~ sP2), inference(cnf_transformation, [], [f208])).
fof(f1607, plain, (~ spl78_139 | spl78_61), inference(avatar_split_clause, [], [f639, f945, f1601])).
fof(f639, plain, ((e0 = op(e0, e0)) | ~ sP3), inference(cnf_transformation, [], [f207])).
fof(f207, plain, ((~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP3), inference(nnf_transformation, [], [f57])).
fof(f1606, plain, (~ spl78_139 | ~ spl78_61), inference(avatar_split_clause, [], [f640, f945, f1601])).
fof(f640, plain, (~ (e0 = op(e0, e0)) | ~ sP3), inference(cnf_transformation, [], [f207])).
fof(f1599, plain, (~ spl78_138 | spl78_57), inference(avatar_split_clause, [], [f635, f928, f1593])).
fof(f635, plain, ((e0 = op(e0, e1)) | ~ sP4), inference(cnf_transformation, [], [f206])).
fof(f206, plain, ((~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | ~ sP4), inference(nnf_transformation, [], [f58])).
fof(f1598, plain, (~ spl78_138 | ~ spl78_62), inference(avatar_split_clause, [], [f636, f949, f1593])).
fof(f636, plain, (~ (op(e0, e0) = e1) | ~ sP4), inference(cnf_transformation, [], [f206])).
fof(f1597, plain, (~ spl78_138 | spl78_62), inference(avatar_split_clause, [], [f637, f949, f1593])).
fof(f637, plain, ((op(e0, e0) = e1) | ~ sP4), inference(cnf_transformation, [], [f206])).
fof(f1589, plain, (~ spl78_137 | spl78_42), inference(avatar_split_clause, [], [f633, f864, f1585])).
fof(f633, plain, ((e1 = op(e1, e1)) | ~ sP5), inference(cnf_transformation, [], [f205])).
fof(f205, plain, ((~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | ~ sP5), inference(nnf_transformation, [], [f59])).
fof(f1588, plain, (~ spl78_137 | ~ spl78_42), inference(avatar_split_clause, [], [f634, f864, f1585])).
fof(f634, plain, (~ (e1 = op(e1, e1)) | ~ sP5), inference(cnf_transformation, [], [f205])).
fof(f1583, plain, (~ spl78_136 | spl78_57), inference(avatar_split_clause, [], [f627, f928, f1577])).
fof(f627, plain, ((e0 = op(e0, e1)) | ~ sP6), inference(cnf_transformation, [], [f204])).
fof(f204, plain, ((~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | ~ sP6), inference(nnf_transformation, [], [f60])).
fof(f1582, plain, (~ spl78_136 | ~ spl78_62), inference(avatar_split_clause, [], [f628, f949, f1577])).
fof(f628, plain, (~ (op(e0, e0) = e1) | ~ sP6), inference(cnf_transformation, [], [f204])).
fof(f1581, plain, (~ spl78_136 | spl78_22), inference(avatar_split_clause, [], [f629, f779, f1577])).
fof(f629, plain, ((e1 = op(e2, e2)) | ~ sP6), inference(cnf_transformation, [], [f204])).
fof(f1580, plain, (~ spl78_136 | ~ spl78_27), inference(avatar_split_clause, [], [f630, f800, f1577])).
fof(f630, plain, (~ (e2 = op(e2, e1)) | ~ sP6), inference(cnf_transformation, [], [f204])).
fof(f1575, plain, (~ spl78_135 | spl78_57), inference(avatar_split_clause, [], [f623, f928, f1569])).
fof(f623, plain, ((e0 = op(e0, e1)) | ~ sP7), inference(cnf_transformation, [], [f203])).
fof(f203, plain, ((~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | ~ sP7), inference(nnf_transformation, [], [f61])).
fof(f1574, plain, (~ spl78_135 | ~ spl78_62), inference(avatar_split_clause, [], [f624, f949, f1569])).
fof(f624, plain, (~ (op(e0, e0) = e1) | ~ sP7), inference(cnf_transformation, [], [f203])).
fof(f1573, plain, (~ spl78_135 | spl78_2), inference(avatar_split_clause, [], [f625, f694, f1569])).
fof(f625, plain, ((e1 = op(e3, e3)) | ~ sP7), inference(cnf_transformation, [], [f203])).
fof(f1572, plain, (~ spl78_135 | ~ spl78_12), inference(avatar_split_clause, [], [f626, f736, f1569])).
fof(f626, plain, (~ (e3 = op(e3, e1)) | ~ sP7), inference(cnf_transformation, [], [f203])).
fof(f1566, plain, (~ spl78_134 | ~ spl78_63), inference(avatar_split_clause, [], [f620, f953, f1561])).
fof(f620, plain, (~ (op(e0, e0) = e2) | ~ sP8), inference(cnf_transformation, [], [f202])).
fof(f202, plain, ((~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | ~ sP8), inference(nnf_transformation, [], [f62])).
fof(f1565, plain, (~ spl78_134 | spl78_63), inference(avatar_split_clause, [], [f621, f953, f1561])).
fof(f621, plain, ((op(e0, e0) = e2) | ~ sP8), inference(cnf_transformation, [], [f202])).
fof(f1559, plain, (~ spl78_133 | spl78_53), inference(avatar_split_clause, [], [f615, f911, f1553])).
fof(f615, plain, ((e0 = op(e0, e2)) | ~ sP9), inference(cnf_transformation, [], [f201])).
fof(f201, plain, ((~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | ~ sP9), inference(nnf_transformation, [], [f63])).
fof(f1558, plain, (~ spl78_133 | ~ spl78_63), inference(avatar_split_clause, [], [f616, f953, f1553])).
fof(f616, plain, (~ (op(e0, e0) = e2) | ~ sP9), inference(cnf_transformation, [], [f201])).
fof(f1557, plain, (~ spl78_133 | spl78_43), inference(avatar_split_clause, [], [f617, f868, f1553])).
fof(f617, plain, ((e2 = op(e1, e1)) | ~ sP9), inference(cnf_transformation, [], [f201])).
fof(f1556, plain, (~ spl78_133 | ~ spl78_38), inference(avatar_split_clause, [], [f618, f847, f1553])).
fof(f618, plain, (~ (e1 = op(e1, e2)) | ~ sP9), inference(cnf_transformation, [], [f201])).
fof(f1549, plain, (~ spl78_132 | spl78_23), inference(avatar_split_clause, [], [f613, f783, f1545])).
fof(f613, plain, ((e2 = op(e2, e2)) | ~ sP10), inference(cnf_transformation, [], [f200])).
fof(f200, plain, ((~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | ~ sP10), inference(nnf_transformation, [], [f64])).
fof(f1548, plain, (~ spl78_132 | ~ spl78_23), inference(avatar_split_clause, [], [f614, f783, f1545])).
fof(f614, plain, (~ (e2 = op(e2, e2)) | ~ sP10), inference(cnf_transformation, [], [f200])).
fof(f1543, plain, (~ spl78_131 | spl78_53), inference(avatar_split_clause, [], [f607, f911, f1537])).
fof(f607, plain, ((e0 = op(e0, e2)) | ~ sP11), inference(cnf_transformation, [], [f199])).
fof(f199, plain, ((~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | ~ sP11), inference(nnf_transformation, [], [f65])).
fof(f1542, plain, (~ spl78_131 | ~ spl78_63), inference(avatar_split_clause, [], [f608, f953, f1537])).
fof(f608, plain, (~ (op(e0, e0) = e2) | ~ sP11), inference(cnf_transformation, [], [f199])).
fof(f1541, plain, (~ spl78_131 | spl78_3), inference(avatar_split_clause, [], [f609, f698, f1537])).
fof(f609, plain, ((e2 = op(e3, e3)) | ~ sP11), inference(cnf_transformation, [], [f199])).
fof(f1540, plain, (~ spl78_131 | ~ spl78_8), inference(avatar_split_clause, [], [f610, f719, f1537])).
fof(f610, plain, (~ (e3 = op(e3, e2)) | ~ sP11), inference(cnf_transformation, [], [f199])).
fof(f1535, plain, (~ spl78_130 | spl78_49), inference(avatar_split_clause, [], [f603, f894, f1529])).
fof(f603, plain, ((e0 = op(e0, e3)) | ~ sP12), inference(cnf_transformation, [], [f198])).
fof(f198, plain, ((~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | ~ sP12), inference(nnf_transformation, [], [f66])).
fof(f1534, plain, (~ spl78_130 | ~ spl78_64), inference(avatar_split_clause, [], [f604, f957, f1529])).
fof(f604, plain, (~ (op(e0, e0) = e3) | ~ sP12), inference(cnf_transformation, [], [f198])).
fof(f1533, plain, (~ spl78_130 | spl78_64), inference(avatar_split_clause, [], [f605, f957, f1529])).
fof(f605, plain, ((op(e0, e0) = e3) | ~ sP12), inference(cnf_transformation, [], [f198])).
fof(f1527, plain, (~ spl78_129 | spl78_49), inference(avatar_split_clause, [], [f599, f894, f1521])).
fof(f599, plain, ((e0 = op(e0, e3)) | ~ sP13), inference(cnf_transformation, [], [f197])).
fof(f197, plain, ((~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | ~ sP13), inference(nnf_transformation, [], [f67])).
fof(f1526, plain, (~ spl78_129 | ~ spl78_64), inference(avatar_split_clause, [], [f600, f957, f1521])).
fof(f600, plain, (~ (op(e0, e0) = e3) | ~ sP13), inference(cnf_transformation, [], [f197])).
fof(f1525, plain, (~ spl78_129 | spl78_44), inference(avatar_split_clause, [], [f601, f872, f1521])).
fof(f601, plain, ((e3 = op(e1, e1)) | ~ sP13), inference(cnf_transformation, [], [f197])).
fof(f1524, plain, (~ spl78_129 | ~ spl78_34), inference(avatar_split_clause, [], [f602, f830, f1521])).
fof(f602, plain, (~ (e1 = op(e1, e3)) | ~ sP13), inference(cnf_transformation, [], [f197])).
fof(f1519, plain, (~ spl78_128 | spl78_49), inference(avatar_split_clause, [], [f595, f894, f1513])).
fof(f595, plain, ((e0 = op(e0, e3)) | ~ sP14), inference(cnf_transformation, [], [f196])).
fof(f196, plain, ((~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | ~ sP14), inference(nnf_transformation, [], [f68])).
fof(f1518, plain, (~ spl78_128 | ~ spl78_64), inference(avatar_split_clause, [], [f596, f957, f1513])).
fof(f596, plain, (~ (op(e0, e0) = e3) | ~ sP14), inference(cnf_transformation, [], [f196])).
fof(f1517, plain, (~ spl78_128 | spl78_24), inference(avatar_split_clause, [], [f597, f787, f1513])).
fof(f597, plain, ((e3 = op(e2, e2)) | ~ sP14), inference(cnf_transformation, [], [f196])).
fof(f1516, plain, (~ spl78_128 | ~ spl78_19), inference(avatar_split_clause, [], [f598, f766, f1513])).
fof(f598, plain, (~ (e2 = op(e2, e3)) | ~ sP14), inference(cnf_transformation, [], [f196])).
fof(f1509, plain, (~ spl78_127 | spl78_4), inference(avatar_split_clause, [], [f593, f702, f1505])).
fof(f593, plain, ((e3 = op(e3, e3)) | ~ sP15), inference(cnf_transformation, [], [f195])).
fof(f195, plain, ((~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | ~ sP15), inference(nnf_transformation, [], [f69])).
fof(f1508, plain, (~ spl78_127 | ~ spl78_4), inference(avatar_split_clause, [], [f594, f702, f1505])).
fof(f594, plain, (~ (e3 = op(e3, e3)) | ~ sP15), inference(cnf_transformation, [], [f195])).
fof(f1503, plain, (~ spl78_126 | spl78_46), inference(avatar_split_clause, [], [f587, f881, f1497])).
fof(f587, plain, ((e1 = op(e1, e0)) | ~ sP16), inference(cnf_transformation, [], [f194])).
fof(f194, plain, ((~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | ~ sP16), inference(nnf_transformation, [], [f70])).
fof(f1501, plain, (~ spl78_126 | spl78_61), inference(avatar_split_clause, [], [f589, f945, f1497])).
fof(f589, plain, ((e0 = op(e0, e0)) | ~ sP16), inference(cnf_transformation, [], [f194])).
fof(f1500, plain, (~ spl78_126 | ~ spl78_61), inference(avatar_split_clause, [], [f590, f945, f1497])).
fof(f590, plain, (~ (e0 = op(e0, e0)) | ~ sP16), inference(cnf_transformation, [], [f194])).
fof(f1495, plain, (~ spl78_125 | spl78_46), inference(avatar_split_clause, [], [f583, f881, f1489])).
fof(f583, plain, ((e1 = op(e1, e0)) | ~ sP17), inference(cnf_transformation, [], [f193])).
fof(f193, plain, ((~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | ~ sP17), inference(nnf_transformation, [], [f71])).
fof(f1494, plain, (~ spl78_125 | ~ spl78_41), inference(avatar_split_clause, [], [f584, f860, f1489])).
fof(f584, plain, (~ (e0 = op(e1, e1)) | ~ sP17), inference(cnf_transformation, [], [f193])).
fof(f1493, plain, (~ spl78_125 | spl78_41), inference(avatar_split_clause, [], [f585, f860, f1489])).
fof(f585, plain, ((e0 = op(e1, e1)) | ~ sP17), inference(cnf_transformation, [], [f193])).
fof(f1487, plain, (~ spl78_124 | spl78_46), inference(avatar_split_clause, [], [f579, f881, f1481])).
fof(f579, plain, ((e1 = op(e1, e0)) | ~ sP18), inference(cnf_transformation, [], [f192])).
fof(f192, plain, ((~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | ~ sP18), inference(nnf_transformation, [], [f72])).
fof(f1486, plain, (~ spl78_124 | ~ spl78_41), inference(avatar_split_clause, [], [f580, f860, f1481])).
fof(f580, plain, (~ (e0 = op(e1, e1)) | ~ sP18), inference(cnf_transformation, [], [f192])).
fof(f1485, plain, (~ spl78_124 | spl78_21), inference(avatar_split_clause, [], [f581, f775, f1481])).
fof(f581, plain, ((e0 = op(e2, e2)) | ~ sP18), inference(cnf_transformation, [], [f192])).
fof(f1484, plain, (~ spl78_124 | ~ spl78_31), inference(avatar_split_clause, [], [f582, f817, f1481])).
fof(f582, plain, (~ (e2 = op(e2, e0)) | ~ sP18), inference(cnf_transformation, [], [f192])).
fof(f1479, plain, (~ spl78_123 | spl78_46), inference(avatar_split_clause, [], [f575, f881, f1473])).
fof(f575, plain, ((e1 = op(e1, e0)) | ~ sP19), inference(cnf_transformation, [], [f191])).
fof(f191, plain, ((~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | ~ sP19), inference(nnf_transformation, [], [f73])).
fof(f1478, plain, (~ spl78_123 | ~ spl78_41), inference(avatar_split_clause, [], [f576, f860, f1473])).
fof(f576, plain, (~ (e0 = op(e1, e1)) | ~ sP19), inference(cnf_transformation, [], [f191])).
fof(f1477, plain, (~ spl78_123 | spl78_1), inference(avatar_split_clause, [], [f577, f690, f1473])).
fof(f577, plain, ((e0 = op(e3, e3)) | ~ sP19), inference(cnf_transformation, [], [f191])).
fof(f1476, plain, (~ spl78_123 | ~ spl78_16), inference(avatar_split_clause, [], [f578, f753, f1473])).
fof(f578, plain, (~ (e3 = op(e3, e0)) | ~ sP19), inference(cnf_transformation, [], [f191])).
fof(f1471, plain, (~ spl78_122 | spl78_42), inference(avatar_split_clause, [], [f571, f864, f1465])).
fof(f571, plain, ((e1 = op(e1, e1)) | ~ sP20), inference(cnf_transformation, [], [f190])).
fof(f190, plain, ((~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP20), inference(nnf_transformation, [], [f74])).
fof(f1470, plain, (~ spl78_122 | ~ spl78_42), inference(avatar_split_clause, [], [f572, f864, f1465])).
fof(f572, plain, (~ (e1 = op(e1, e1)) | ~ sP20), inference(cnf_transformation, [], [f190])).
fof(f1463, plain, (~ spl78_121 | spl78_42), inference(avatar_split_clause, [], [f567, f864, f1457])).
fof(f567, plain, ((e1 = op(e1, e1)) | ~ sP21), inference(cnf_transformation, [], [f189])).
fof(f189, plain, ((~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP21), inference(nnf_transformation, [], [f75])).
fof(f1462, plain, (~ spl78_121 | ~ spl78_42), inference(avatar_split_clause, [], [f568, f864, f1457])).
fof(f568, plain, (~ (e1 = op(e1, e1)) | ~ sP21), inference(cnf_transformation, [], [f189])).
fof(f1455, plain, (~ spl78_120 | spl78_42), inference(avatar_split_clause, [], [f563, f864, f1449])).
fof(f563, plain, ((e1 = op(e1, e1)) | ~ sP22), inference(cnf_transformation, [], [f188])).
fof(f188, plain, ((~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP22), inference(nnf_transformation, [], [f76])).
fof(f1454, plain, (~ spl78_120 | ~ spl78_42), inference(avatar_split_clause, [], [f564, f864, f1449])).
fof(f564, plain, (~ (e1 = op(e1, e1)) | ~ sP22), inference(cnf_transformation, [], [f188])).
fof(f1447, plain, (~ spl78_119 | spl78_42), inference(avatar_split_clause, [], [f559, f864, f1441])).
fof(f559, plain, ((e1 = op(e1, e1)) | ~ sP23), inference(cnf_transformation, [], [f187])).
fof(f187, plain, ((~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP23), inference(nnf_transformation, [], [f77])).
fof(f1446, plain, (~ spl78_119 | ~ spl78_42), inference(avatar_split_clause, [], [f560, f864, f1441])).
fof(f560, plain, (~ (e1 = op(e1, e1)) | ~ sP23), inference(cnf_transformation, [], [f187])).
fof(f1439, plain, (~ spl78_118 | spl78_38), inference(avatar_split_clause, [], [f555, f847, f1433])).
fof(f555, plain, ((e1 = op(e1, e2)) | ~ sP24), inference(cnf_transformation, [], [f186])).
fof(f186, plain, ((~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | ~ sP24), inference(nnf_transformation, [], [f78])).
fof(f1438, plain, (~ spl78_118 | ~ spl78_43), inference(avatar_split_clause, [], [f556, f868, f1433])).
fof(f556, plain, (~ (e2 = op(e1, e1)) | ~ sP24), inference(cnf_transformation, [], [f186])).
fof(f1437, plain, (~ spl78_118 | spl78_63), inference(avatar_split_clause, [], [f557, f953, f1433])).
fof(f557, plain, ((op(e0, e0) = e2) | ~ sP24), inference(cnf_transformation, [], [f186])).
fof(f1436, plain, (~ spl78_118 | ~ spl78_53), inference(avatar_split_clause, [], [f558, f911, f1433])).
fof(f558, plain, (~ (e0 = op(e0, e2)) | ~ sP24), inference(cnf_transformation, [], [f186])).
fof(f1430, plain, (~ spl78_117 | ~ spl78_43), inference(avatar_split_clause, [], [f552, f868, f1425])).
fof(f552, plain, (~ (e2 = op(e1, e1)) | ~ sP25), inference(cnf_transformation, [], [f185])).
fof(f185, plain, ((~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | ~ sP25), inference(nnf_transformation, [], [f79])).
fof(f1429, plain, (~ spl78_117 | spl78_43), inference(avatar_split_clause, [], [f553, f868, f1425])).
fof(f553, plain, ((e2 = op(e1, e1)) | ~ sP25), inference(cnf_transformation, [], [f185])).
fof(f1421, plain, (~ spl78_116 | spl78_23), inference(avatar_split_clause, [], [f549, f783, f1417])).
fof(f549, plain, ((e2 = op(e2, e2)) | ~ sP26), inference(cnf_transformation, [], [f184])).
fof(f184, plain, ((~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | ~ sP26), inference(nnf_transformation, [], [f80])).
fof(f1420, plain, (~ spl78_116 | ~ spl78_23), inference(avatar_split_clause, [], [f550, f783, f1417])).
fof(f550, plain, (~ (e2 = op(e2, e2)) | ~ sP26), inference(cnf_transformation, [], [f184])).
fof(f1415, plain, (~ spl78_115 | spl78_38), inference(avatar_split_clause, [], [f543, f847, f1409])).
fof(f543, plain, ((e1 = op(e1, e2)) | ~ sP27), inference(cnf_transformation, [], [f183])).
fof(f183, plain, ((~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | ~ sP27), inference(nnf_transformation, [], [f81])).
fof(f1414, plain, (~ spl78_115 | ~ spl78_43), inference(avatar_split_clause, [], [f544, f868, f1409])).
fof(f544, plain, (~ (e2 = op(e1, e1)) | ~ sP27), inference(cnf_transformation, [], [f183])).
fof(f1413, plain, (~ spl78_115 | spl78_3), inference(avatar_split_clause, [], [f545, f698, f1409])).
fof(f545, plain, ((e2 = op(e3, e3)) | ~ sP27), inference(cnf_transformation, [], [f183])).
fof(f1412, plain, (~ spl78_115 | ~ spl78_8), inference(avatar_split_clause, [], [f546, f719, f1409])).
fof(f546, plain, (~ (e3 = op(e3, e2)) | ~ sP27), inference(cnf_transformation, [], [f183])).
fof(f1407, plain, (~ spl78_114 | spl78_34), inference(avatar_split_clause, [], [f539, f830, f1401])).
fof(f539, plain, ((e1 = op(e1, e3)) | ~ sP28), inference(cnf_transformation, [], [f182])).
fof(f182, plain, ((~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | ~ sP28), inference(nnf_transformation, [], [f82])).
fof(f1406, plain, (~ spl78_114 | ~ spl78_44), inference(avatar_split_clause, [], [f540, f872, f1401])).
fof(f540, plain, (~ (e3 = op(e1, e1)) | ~ sP28), inference(cnf_transformation, [], [f182])).
fof(f1405, plain, (~ spl78_114 | spl78_64), inference(avatar_split_clause, [], [f541, f957, f1401])).
fof(f541, plain, ((op(e0, e0) = e3) | ~ sP28), inference(cnf_transformation, [], [f182])).
fof(f1404, plain, (~ spl78_114 | ~ spl78_49), inference(avatar_split_clause, [], [f542, f894, f1401])).
fof(f542, plain, (~ (e0 = op(e0, e3)) | ~ sP28), inference(cnf_transformation, [], [f182])).
fof(f1398, plain, (~ spl78_113 | ~ spl78_44), inference(avatar_split_clause, [], [f536, f872, f1393])).
fof(f536, plain, (~ (e3 = op(e1, e1)) | ~ sP29), inference(cnf_transformation, [], [f181])).
fof(f181, plain, ((~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | ~ sP29), inference(nnf_transformation, [], [f83])).
fof(f1397, plain, (~ spl78_113 | spl78_44), inference(avatar_split_clause, [], [f537, f872, f1393])).
fof(f537, plain, ((e3 = op(e1, e1)) | ~ sP29), inference(cnf_transformation, [], [f181])).
fof(f1391, plain, (~ spl78_112 | spl78_34), inference(avatar_split_clause, [], [f531, f830, f1385])).
fof(f531, plain, ((e1 = op(e1, e3)) | ~ sP30), inference(cnf_transformation, [], [f180])).
fof(f180, plain, ((~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | ~ sP30), inference(nnf_transformation, [], [f84])).
fof(f1390, plain, (~ spl78_112 | ~ spl78_44), inference(avatar_split_clause, [], [f532, f872, f1385])).
fof(f532, plain, (~ (e3 = op(e1, e1)) | ~ sP30), inference(cnf_transformation, [], [f180])).
fof(f1389, plain, (~ spl78_112 | spl78_24), inference(avatar_split_clause, [], [f533, f787, f1385])).
fof(f533, plain, ((e3 = op(e2, e2)) | ~ sP30), inference(cnf_transformation, [], [f180])).
fof(f1388, plain, (~ spl78_112 | ~ spl78_19), inference(avatar_split_clause, [], [f534, f766, f1385])).
fof(f534, plain, (~ (e2 = op(e2, e3)) | ~ sP30), inference(cnf_transformation, [], [f180])).
fof(f1381, plain, (~ spl78_111 | spl78_4), inference(avatar_split_clause, [], [f529, f702, f1377])).
fof(f529, plain, ((e3 = op(e3, e3)) | ~ sP31), inference(cnf_transformation, [], [f179])).
fof(f179, plain, ((~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | ~ sP31), inference(nnf_transformation, [], [f85])).
fof(f1380, plain, (~ spl78_111 | ~ spl78_4), inference(avatar_split_clause, [], [f530, f702, f1377])).
fof(f530, plain, (~ (e3 = op(e3, e3)) | ~ sP31), inference(cnf_transformation, [], [f179])).
fof(f1375, plain, (~ spl78_110 | spl78_31), inference(avatar_split_clause, [], [f523, f817, f1369])).
fof(f523, plain, ((e2 = op(e2, e0)) | ~ sP32), inference(cnf_transformation, [], [f178])).
fof(f178, plain, ((~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | ~ sP32), inference(nnf_transformation, [], [f86])).
fof(f1373, plain, (~ spl78_110 | spl78_61), inference(avatar_split_clause, [], [f525, f945, f1369])).
fof(f525, plain, ((e0 = op(e0, e0)) | ~ sP32), inference(cnf_transformation, [], [f178])).
fof(f1372, plain, (~ spl78_110 | ~ spl78_61), inference(avatar_split_clause, [], [f526, f945, f1369])).
fof(f526, plain, (~ (e0 = op(e0, e0)) | ~ sP32), inference(cnf_transformation, [], [f178])).
fof(f1367, plain, (~ spl78_109 | spl78_31), inference(avatar_split_clause, [], [f519, f817, f1361])).
fof(f519, plain, ((e2 = op(e2, e0)) | ~ sP33), inference(cnf_transformation, [], [f177])).
fof(f177, plain, ((~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | ~ sP33), inference(nnf_transformation, [], [f87])).
fof(f1366, plain, (~ spl78_109 | ~ spl78_21), inference(avatar_split_clause, [], [f520, f775, f1361])).
fof(f520, plain, (~ (e0 = op(e2, e2)) | ~ sP33), inference(cnf_transformation, [], [f177])).
fof(f1365, plain, (~ spl78_109 | spl78_41), inference(avatar_split_clause, [], [f521, f860, f1361])).
fof(f521, plain, ((e0 = op(e1, e1)) | ~ sP33), inference(cnf_transformation, [], [f177])).
fof(f1364, plain, (~ spl78_109 | ~ spl78_46), inference(avatar_split_clause, [], [f522, f881, f1361])).
fof(f522, plain, (~ (e1 = op(e1, e0)) | ~ sP33), inference(cnf_transformation, [], [f177])).
fof(f1359, plain, (~ spl78_108 | spl78_31), inference(avatar_split_clause, [], [f515, f817, f1353])).
fof(f515, plain, ((e2 = op(e2, e0)) | ~ sP34), inference(cnf_transformation, [], [f176])).
fof(f176, plain, ((~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | ~ sP34), inference(nnf_transformation, [], [f88])).
fof(f1358, plain, (~ spl78_108 | ~ spl78_21), inference(avatar_split_clause, [], [f516, f775, f1353])).
fof(f516, plain, (~ (e0 = op(e2, e2)) | ~ sP34), inference(cnf_transformation, [], [f176])).
fof(f1357, plain, (~ spl78_108 | spl78_21), inference(avatar_split_clause, [], [f517, f775, f1353])).
fof(f517, plain, ((e0 = op(e2, e2)) | ~ sP34), inference(cnf_transformation, [], [f176])).
fof(f1351, plain, (~ spl78_107 | spl78_31), inference(avatar_split_clause, [], [f511, f817, f1345])).
fof(f511, plain, ((e2 = op(e2, e0)) | ~ sP35), inference(cnf_transformation, [], [f175])).
fof(f175, plain, ((~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | ~ sP35), inference(nnf_transformation, [], [f89])).
fof(f1350, plain, (~ spl78_107 | ~ spl78_21), inference(avatar_split_clause, [], [f512, f775, f1345])).
fof(f512, plain, (~ (e0 = op(e2, e2)) | ~ sP35), inference(cnf_transformation, [], [f175])).
fof(f1349, plain, (~ spl78_107 | spl78_1), inference(avatar_split_clause, [], [f513, f690, f1345])).
fof(f513, plain, ((e0 = op(e3, e3)) | ~ sP35), inference(cnf_transformation, [], [f175])).
fof(f1348, plain, (~ spl78_107 | ~ spl78_16), inference(avatar_split_clause, [], [f514, f753, f1345])).
fof(f514, plain, (~ (e3 = op(e3, e0)) | ~ sP35), inference(cnf_transformation, [], [f175])).
fof(f1343, plain, (~ spl78_106 | spl78_27), inference(avatar_split_clause, [], [f507, f800, f1337])).
fof(f507, plain, ((e2 = op(e2, e1)) | ~ sP36), inference(cnf_transformation, [], [f174])).
fof(f174, plain, ((~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | ~ sP36), inference(nnf_transformation, [], [f90])).
fof(f1342, plain, (~ spl78_106 | ~ spl78_22), inference(avatar_split_clause, [], [f508, f779, f1337])).
fof(f508, plain, (~ (e1 = op(e2, e2)) | ~ sP36), inference(cnf_transformation, [], [f174])).
fof(f1341, plain, (~ spl78_106 | spl78_62), inference(avatar_split_clause, [], [f509, f949, f1337])).
fof(f509, plain, ((op(e0, e0) = e1) | ~ sP36), inference(cnf_transformation, [], [f174])).
fof(f1340, plain, (~ spl78_106 | ~ spl78_57), inference(avatar_split_clause, [], [f510, f928, f1337])).
fof(f510, plain, (~ (e0 = op(e0, e1)) | ~ sP36), inference(cnf_transformation, [], [f174])).
fof(f1333, plain, (~ spl78_105 | spl78_42), inference(avatar_split_clause, [], [f505, f864, f1329])).
fof(f505, plain, ((e1 = op(e1, e1)) | ~ sP37), inference(cnf_transformation, [], [f173])).
fof(f173, plain, ((~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | ~ sP37), inference(nnf_transformation, [], [f91])).
fof(f1332, plain, (~ spl78_105 | ~ spl78_42), inference(avatar_split_clause, [], [f506, f864, f1329])).
fof(f506, plain, (~ (e1 = op(e1, e1)) | ~ sP37), inference(cnf_transformation, [], [f173])).
fof(f1326, plain, (~ spl78_104 | ~ spl78_22), inference(avatar_split_clause, [], [f500, f779, f1321])).
fof(f500, plain, (~ (e1 = op(e2, e2)) | ~ sP38), inference(cnf_transformation, [], [f172])).
fof(f172, plain, ((~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | ~ sP38), inference(nnf_transformation, [], [f92])).
fof(f1325, plain, (~ spl78_104 | spl78_22), inference(avatar_split_clause, [], [f501, f779, f1321])).
fof(f501, plain, ((e1 = op(e2, e2)) | ~ sP38), inference(cnf_transformation, [], [f172])).
fof(f1319, plain, (~ spl78_103 | spl78_27), inference(avatar_split_clause, [], [f495, f800, f1313])).
fof(f495, plain, ((e2 = op(e2, e1)) | ~ sP39), inference(cnf_transformation, [], [f171])).
fof(f171, plain, ((~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | ~ sP39), inference(nnf_transformation, [], [f93])).
fof(f1318, plain, (~ spl78_103 | ~ spl78_22), inference(avatar_split_clause, [], [f496, f779, f1313])).
fof(f496, plain, (~ (e1 = op(e2, e2)) | ~ sP39), inference(cnf_transformation, [], [f171])).
fof(f1317, plain, (~ spl78_103 | spl78_2), inference(avatar_split_clause, [], [f497, f694, f1313])).
fof(f497, plain, ((e1 = op(e3, e3)) | ~ sP39), inference(cnf_transformation, [], [f171])).
fof(f1316, plain, (~ spl78_103 | ~ spl78_12), inference(avatar_split_clause, [], [f498, f736, f1313])).
fof(f498, plain, (~ (e3 = op(e3, e1)) | ~ sP39), inference(cnf_transformation, [], [f171])).
fof(f1311, plain, (~ spl78_102 | spl78_23), inference(avatar_split_clause, [], [f491, f783, f1305])).
fof(f491, plain, ((e2 = op(e2, e2)) | ~ sP40), inference(cnf_transformation, [], [f170])).
fof(f170, plain, ((~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP40), inference(nnf_transformation, [], [f94])).
fof(f1310, plain, (~ spl78_102 | ~ spl78_23), inference(avatar_split_clause, [], [f492, f783, f1305])).
fof(f492, plain, (~ (e2 = op(e2, e2)) | ~ sP40), inference(cnf_transformation, [], [f170])).
fof(f1303, plain, (~ spl78_101 | spl78_23), inference(avatar_split_clause, [], [f487, f783, f1297])).
fof(f487, plain, ((e2 = op(e2, e2)) | ~ sP41), inference(cnf_transformation, [], [f169])).
fof(f169, plain, ((~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP41), inference(nnf_transformation, [], [f95])).
fof(f1302, plain, (~ spl78_101 | ~ spl78_23), inference(avatar_split_clause, [], [f488, f783, f1297])).
fof(f488, plain, (~ (e2 = op(e2, e2)) | ~ sP41), inference(cnf_transformation, [], [f169])).
fof(f1301, plain, (~ spl78_101 | spl78_43), inference(avatar_split_clause, [], [f489, f868, f1297])).
fof(f489, plain, ((e2 = op(e1, e1)) | ~ sP41), inference(cnf_transformation, [], [f169])).
fof(f1295, plain, (~ spl78_100 | spl78_23), inference(avatar_split_clause, [], [f483, f783, f1289])).
fof(f483, plain, ((e2 = op(e2, e2)) | ~ sP42), inference(cnf_transformation, [], [f168])).
fof(f168, plain, ((~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP42), inference(nnf_transformation, [], [f96])).
fof(f1294, plain, (~ spl78_100 | ~ spl78_23), inference(avatar_split_clause, [], [f484, f783, f1289])).
fof(f484, plain, (~ (e2 = op(e2, e2)) | ~ sP42), inference(cnf_transformation, [], [f168])).
fof(f1287, plain, (~ spl78_99 | spl78_23), inference(avatar_split_clause, [], [f479, f783, f1281])).
fof(f479, plain, ((e2 = op(e2, e2)) | ~ sP43), inference(cnf_transformation, [], [f167])).
fof(f167, plain, ((~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP43), inference(nnf_transformation, [], [f97])).
fof(f1286, plain, (~ spl78_99 | ~ spl78_23), inference(avatar_split_clause, [], [f480, f783, f1281])).
fof(f480, plain, (~ (e2 = op(e2, e2)) | ~ sP43), inference(cnf_transformation, [], [f167])).
fof(f1279, plain, (~ spl78_98 | spl78_19), inference(avatar_split_clause, [], [f475, f766, f1273])).
fof(f475, plain, ((e2 = op(e2, e3)) | ~ sP44), inference(cnf_transformation, [], [f166])).
fof(f166, plain, ((~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | ~ sP44), inference(nnf_transformation, [], [f98])).
fof(f1278, plain, (~ spl78_98 | ~ spl78_24), inference(avatar_split_clause, [], [f476, f787, f1273])).
fof(f476, plain, (~ (e3 = op(e2, e2)) | ~ sP44), inference(cnf_transformation, [], [f166])).
fof(f1277, plain, (~ spl78_98 | spl78_64), inference(avatar_split_clause, [], [f477, f957, f1273])).
fof(f477, plain, ((op(e0, e0) = e3) | ~ sP44), inference(cnf_transformation, [], [f166])).
fof(f1276, plain, (~ spl78_98 | ~ spl78_49), inference(avatar_split_clause, [], [f478, f894, f1273])).
fof(f478, plain, (~ (e0 = op(e0, e3)) | ~ sP44), inference(cnf_transformation, [], [f166])).
fof(f1271, plain, (~ spl78_97 | spl78_19), inference(avatar_split_clause, [], [f471, f766, f1265])).
fof(f471, plain, ((e2 = op(e2, e3)) | ~ sP45), inference(cnf_transformation, [], [f165])).
fof(f165, plain, ((~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | ~ sP45), inference(nnf_transformation, [], [f99])).
fof(f1270, plain, (~ spl78_97 | ~ spl78_24), inference(avatar_split_clause, [], [f472, f787, f1265])).
fof(f472, plain, (~ (e3 = op(e2, e2)) | ~ sP45), inference(cnf_transformation, [], [f165])).
fof(f1269, plain, (~ spl78_97 | spl78_44), inference(avatar_split_clause, [], [f473, f872, f1265])).
fof(f473, plain, ((e3 = op(e1, e1)) | ~ sP45), inference(cnf_transformation, [], [f165])).
fof(f1268, plain, (~ spl78_97 | ~ spl78_34), inference(avatar_split_clause, [], [f474, f830, f1265])).
fof(f474, plain, (~ (e1 = op(e1, e3)) | ~ sP45), inference(cnf_transformation, [], [f165])).
fof(f1262, plain, (~ spl78_96 | ~ spl78_24), inference(avatar_split_clause, [], [f468, f787, f1257])).
fof(f468, plain, (~ (e3 = op(e2, e2)) | ~ sP46), inference(cnf_transformation, [], [f164])).
fof(f164, plain, ((~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | ~ sP46), inference(nnf_transformation, [], [f100])).
fof(f1261, plain, (~ spl78_96 | spl78_24), inference(avatar_split_clause, [], [f469, f787, f1257])).
fof(f469, plain, ((e3 = op(e2, e2)) | ~ sP46), inference(cnf_transformation, [], [f164])).
fof(f1253, plain, (~ spl78_95 | spl78_4), inference(avatar_split_clause, [], [f465, f702, f1249])).
fof(f465, plain, ((e3 = op(e3, e3)) | ~ sP47), inference(cnf_transformation, [], [f163])).
fof(f163, plain, ((~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | ~ sP47), inference(nnf_transformation, [], [f101])).
fof(f1252, plain, (~ spl78_95 | ~ spl78_4), inference(avatar_split_clause, [], [f466, f702, f1249])).
fof(f466, plain, (~ (e3 = op(e3, e3)) | ~ sP47), inference(cnf_transformation, [], [f163])).
fof(f1245, plain, (~ spl78_94 | spl78_61), inference(avatar_split_clause, [], [f461, f945, f1241])).
fof(f461, plain, ((e0 = op(e0, e0)) | ~ sP48), inference(cnf_transformation, [], [f162])).
fof(f162, plain, ((~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | ~ sP48), inference(nnf_transformation, [], [f102])).
fof(f1244, plain, (~ spl78_94 | ~ spl78_61), inference(avatar_split_clause, [], [f462, f945, f1241])).
fof(f462, plain, (~ (e0 = op(e0, e0)) | ~ sP48), inference(cnf_transformation, [], [f162])).
fof(f1239, plain, (~ spl78_93 | spl78_16), inference(avatar_split_clause, [], [f455, f753, f1233])).
fof(f455, plain, ((e3 = op(e3, e0)) | ~ sP49), inference(cnf_transformation, [], [f161])).
fof(f161, plain, ((~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | ~ sP49), inference(nnf_transformation, [], [f103])).
fof(f1238, plain, (~ spl78_93 | ~ spl78_1), inference(avatar_split_clause, [], [f456, f690, f1233])).
fof(f456, plain, (~ (e0 = op(e3, e3)) | ~ sP49), inference(cnf_transformation, [], [f161])).
fof(f1237, plain, (~ spl78_93 | spl78_41), inference(avatar_split_clause, [], [f457, f860, f1233])).
fof(f457, plain, ((e0 = op(e1, e1)) | ~ sP49), inference(cnf_transformation, [], [f161])).
fof(f1236, plain, (~ spl78_93 | ~ spl78_46), inference(avatar_split_clause, [], [f458, f881, f1233])).
fof(f458, plain, (~ (e1 = op(e1, e0)) | ~ sP49), inference(cnf_transformation, [], [f161])).
fof(f1231, plain, (~ spl78_92 | spl78_16), inference(avatar_split_clause, [], [f451, f753, f1225])).
fof(f451, plain, ((e3 = op(e3, e0)) | ~ sP50), inference(cnf_transformation, [], [f160])).
fof(f160, plain, ((~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | ~ sP50), inference(nnf_transformation, [], [f104])).
fof(f1230, plain, (~ spl78_92 | ~ spl78_1), inference(avatar_split_clause, [], [f452, f690, f1225])).
fof(f452, plain, (~ (e0 = op(e3, e3)) | ~ sP50), inference(cnf_transformation, [], [f160])).
fof(f1229, plain, (~ spl78_92 | spl78_21), inference(avatar_split_clause, [], [f453, f775, f1225])).
fof(f453, plain, ((e0 = op(e2, e2)) | ~ sP50), inference(cnf_transformation, [], [f160])).
fof(f1228, plain, (~ spl78_92 | ~ spl78_31), inference(avatar_split_clause, [], [f454, f817, f1225])).
fof(f454, plain, (~ (e2 = op(e2, e0)) | ~ sP50), inference(cnf_transformation, [], [f160])).
fof(f1223, plain, (~ spl78_91 | spl78_16), inference(avatar_split_clause, [], [f447, f753, f1217])).
fof(f447, plain, ((e3 = op(e3, e0)) | ~ sP51), inference(cnf_transformation, [], [f159])).
fof(f159, plain, ((~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | ~ sP51), inference(nnf_transformation, [], [f105])).
fof(f1222, plain, (~ spl78_91 | ~ spl78_1), inference(avatar_split_clause, [], [f448, f690, f1217])).
fof(f448, plain, (~ (e0 = op(e3, e3)) | ~ sP51), inference(cnf_transformation, [], [f159])).
fof(f1221, plain, (~ spl78_91 | spl78_1), inference(avatar_split_clause, [], [f449, f690, f1217])).
fof(f449, plain, ((e0 = op(e3, e3)) | ~ sP51), inference(cnf_transformation, [], [f159])).
fof(f1215, plain, (~ spl78_90 | spl78_12), inference(avatar_split_clause, [], [f443, f736, f1209])).
fof(f443, plain, ((e3 = op(e3, e1)) | ~ sP52), inference(cnf_transformation, [], [f158])).
fof(f158, plain, ((~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | ~ sP52), inference(nnf_transformation, [], [f106])).
fof(f1214, plain, (~ spl78_90 | ~ spl78_2), inference(avatar_split_clause, [], [f444, f694, f1209])).
fof(f444, plain, (~ (e1 = op(e3, e3)) | ~ sP52), inference(cnf_transformation, [], [f158])).
fof(f1213, plain, (~ spl78_90 | spl78_62), inference(avatar_split_clause, [], [f445, f949, f1209])).
fof(f445, plain, ((op(e0, e0) = e1) | ~ sP52), inference(cnf_transformation, [], [f158])).
fof(f1212, plain, (~ spl78_90 | ~ spl78_57), inference(avatar_split_clause, [], [f446, f928, f1209])).
fof(f446, plain, (~ (e0 = op(e0, e1)) | ~ sP52), inference(cnf_transformation, [], [f158])).
fof(f1205, plain, (~ spl78_89 | spl78_42), inference(avatar_split_clause, [], [f441, f864, f1201])).
fof(f441, plain, ((e1 = op(e1, e1)) | ~ sP53), inference(cnf_transformation, [], [f157])).
fof(f157, plain, ((~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | ~ sP53), inference(nnf_transformation, [], [f107])).
fof(f1204, plain, (~ spl78_89 | ~ spl78_42), inference(avatar_split_clause, [], [f442, f864, f1201])).
fof(f442, plain, (~ (e1 = op(e1, e1)) | ~ sP53), inference(cnf_transformation, [], [f157])).
fof(f1199, plain, (~ spl78_88 | spl78_12), inference(avatar_split_clause, [], [f435, f736, f1193])).
fof(f435, plain, ((e3 = op(e3, e1)) | ~ sP54), inference(cnf_transformation, [], [f156])).
fof(f156, plain, ((~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | ~ sP54), inference(nnf_transformation, [], [f108])).
fof(f1198, plain, (~ spl78_88 | ~ spl78_2), inference(avatar_split_clause, [], [f436, f694, f1193])).
fof(f436, plain, (~ (e1 = op(e3, e3)) | ~ sP54), inference(cnf_transformation, [], [f156])).
fof(f1197, plain, (~ spl78_88 | spl78_22), inference(avatar_split_clause, [], [f437, f779, f1193])).
fof(f437, plain, ((e1 = op(e2, e2)) | ~ sP54), inference(cnf_transformation, [], [f156])).
fof(f1196, plain, (~ spl78_88 | ~ spl78_27), inference(avatar_split_clause, [], [f438, f800, f1193])).
fof(f438, plain, (~ (e2 = op(e2, e1)) | ~ sP54), inference(cnf_transformation, [], [f156])).
fof(f1190, plain, (~ spl78_87 | ~ spl78_2), inference(avatar_split_clause, [], [f432, f694, f1185])).
fof(f432, plain, (~ (e1 = op(e3, e3)) | ~ sP55), inference(cnf_transformation, [], [f155])).
fof(f155, plain, ((~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | ~ sP55), inference(nnf_transformation, [], [f109])).
fof(f1189, plain, (~ spl78_87 | spl78_2), inference(avatar_split_clause, [], [f433, f694, f1185])).
fof(f433, plain, ((e1 = op(e3, e3)) | ~ sP55), inference(cnf_transformation, [], [f155])).
fof(f1183, plain, (~ spl78_86 | spl78_8), inference(avatar_split_clause, [], [f427, f719, f1177])).
fof(f427, plain, ((e3 = op(e3, e2)) | ~ sP56), inference(cnf_transformation, [], [f154])).
fof(f154, plain, ((~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | ~ sP56), inference(nnf_transformation, [], [f110])).
fof(f1182, plain, (~ spl78_86 | ~ spl78_3), inference(avatar_split_clause, [], [f428, f698, f1177])).
fof(f428, plain, (~ (e2 = op(e3, e3)) | ~ sP56), inference(cnf_transformation, [], [f154])).
fof(f1181, plain, (~ spl78_86 | spl78_63), inference(avatar_split_clause, [], [f429, f953, f1177])).
fof(f429, plain, ((op(e0, e0) = e2) | ~ sP56), inference(cnf_transformation, [], [f154])).
fof(f1180, plain, (~ spl78_86 | ~ spl78_53), inference(avatar_split_clause, [], [f430, f911, f1177])).
fof(f430, plain, (~ (e0 = op(e0, e2)) | ~ sP56), inference(cnf_transformation, [], [f154])).
fof(f1175, plain, (~ spl78_85 | spl78_8), inference(avatar_split_clause, [], [f423, f719, f1169])).
fof(f423, plain, ((e3 = op(e3, e2)) | ~ sP57), inference(cnf_transformation, [], [f153])).
fof(f153, plain, ((~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | ~ sP57), inference(nnf_transformation, [], [f111])).
fof(f1174, plain, (~ spl78_85 | ~ spl78_3), inference(avatar_split_clause, [], [f424, f698, f1169])).
fof(f424, plain, (~ (e2 = op(e3, e3)) | ~ sP57), inference(cnf_transformation, [], [f153])).
fof(f1173, plain, (~ spl78_85 | spl78_43), inference(avatar_split_clause, [], [f425, f868, f1169])).
fof(f425, plain, ((e2 = op(e1, e1)) | ~ sP57), inference(cnf_transformation, [], [f153])).
fof(f1172, plain, (~ spl78_85 | ~ spl78_38), inference(avatar_split_clause, [], [f426, f847, f1169])).
fof(f426, plain, (~ (e1 = op(e1, e2)) | ~ sP57), inference(cnf_transformation, [], [f153])).
fof(f1167, plain, (~ spl78_84 | spl78_8), inference(avatar_split_clause, [], [f419, f719, f1161])).
fof(f419, plain, ((e3 = op(e3, e2)) | ~ sP58), inference(cnf_transformation, [], [f152])).
fof(f152, plain, ((~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | ~ sP58), inference(nnf_transformation, [], [f112])).
fof(f1165, plain, (~ spl78_84 | spl78_23), inference(avatar_split_clause, [], [f421, f783, f1161])).
fof(f421, plain, ((e2 = op(e2, e2)) | ~ sP58), inference(cnf_transformation, [], [f152])).
fof(f1164, plain, (~ spl78_84 | ~ spl78_23), inference(avatar_split_clause, [], [f422, f783, f1161])).
fof(f422, plain, (~ (e2 = op(e2, e2)) | ~ sP58), inference(cnf_transformation, [], [f152])).
fof(f1159, plain, (~ spl78_83 | spl78_8), inference(avatar_split_clause, [], [f415, f719, f1153])).
fof(f415, plain, ((e3 = op(e3, e2)) | ~ sP59), inference(cnf_transformation, [], [f151])).
fof(f151, plain, ((~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | ~ sP59), inference(nnf_transformation, [], [f113])).
fof(f1158, plain, (~ spl78_83 | ~ spl78_3), inference(avatar_split_clause, [], [f416, f698, f1153])).
fof(f416, plain, (~ (e2 = op(e3, e3)) | ~ sP59), inference(cnf_transformation, [], [f151])).
fof(f1157, plain, (~ spl78_83 | spl78_3), inference(avatar_split_clause, [], [f417, f698, f1153])).
fof(f417, plain, ((e2 = op(e3, e3)) | ~ sP59), inference(cnf_transformation, [], [f151])).
fof(f1151, plain, (~ spl78_82 | spl78_4), inference(avatar_split_clause, [], [f411, f702, f1145])).
fof(f411, plain, ((e3 = op(e3, e3)) | ~ sP60), inference(cnf_transformation, [], [f150])).
fof(f150, plain, ((~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ~ sP60), inference(nnf_transformation, [], [f114])).
fof(f1150, plain, (~ spl78_82 | ~ spl78_4), inference(avatar_split_clause, [], [f412, f702, f1145])).
fof(f412, plain, (~ (e3 = op(e3, e3)) | ~ sP60), inference(cnf_transformation, [], [f150])).
fof(f1149, plain, (~ spl78_82 | spl78_64), inference(avatar_split_clause, [], [f413, f957, f1145])).
fof(f413, plain, ((op(e0, e0) = e3) | ~ sP60), inference(cnf_transformation, [], [f150])).
fof(f1143, plain, (~ spl78_81 | spl78_4), inference(avatar_split_clause, [], [f407, f702, f1137])).
fof(f407, plain, ((e3 = op(e3, e3)) | ~ sP61), inference(cnf_transformation, [], [f149])).
fof(f149, plain, ((~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ~ sP61), inference(nnf_transformation, [], [f115])).
fof(f1142, plain, (~ spl78_81 | ~ spl78_4), inference(avatar_split_clause, [], [f408, f702, f1137])).
fof(f408, plain, (~ (e3 = op(e3, e3)) | ~ sP61), inference(cnf_transformation, [], [f149])).
fof(f1135, plain, (~ spl78_80 | spl78_4), inference(avatar_split_clause, [], [f403, f702, f1129])).
fof(f403, plain, ((e3 = op(e3, e3)) | ~ sP62), inference(cnf_transformation, [], [f148])).
fof(f148, plain, ((~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ~ sP62), inference(nnf_transformation, [], [f116])).
fof(f1134, plain, (~ spl78_80 | ~ spl78_4), inference(avatar_split_clause, [], [f404, f702, f1129])).
fof(f404, plain, (~ (e3 = op(e3, e3)) | ~ sP62), inference(cnf_transformation, [], [f148])).
fof(f1127, plain, (~ spl78_79 | spl78_61), inference(avatar_split_clause, [], [f397, f945, f1120])).
fof(f397, plain, ((e0 = op(e0, e0)) | ~ sP63), inference(cnf_transformation, [], [f147])).
fof(f147, plain, ((((e0 = op(e3, e3)) | ~ (e3 = op(e3, e0))) & ((e0 = op(e2, e2)) | ~ (e2 = op(e2, e0))) & ((e0 = op(e1, e1)) | ~ (e1 = op(e1, e0))) & ((e0 = op(e0, e0)) | ~ (e0 = op(e0, e0))) & (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP63), inference(nnf_transformation, [], [f117])).
fof(f1125, plain, (~ spl78_79 | ~ spl78_46 | spl78_41), inference(avatar_split_clause, [], [f400, f860, f881, f1120])).
fof(f400, plain, ((e0 = op(e1, e1)) | ~ (e1 = op(e1, e0)) | ~ sP63), inference(cnf_transformation, [], [f147])).
fof(f1124, plain, (~ spl78_79 | ~ spl78_31 | spl78_21), inference(avatar_split_clause, [], [f401, f775, f817, f1120])).
fof(f401, plain, ((e0 = op(e2, e2)) | ~ (e2 = op(e2, e0)) | ~ sP63), inference(cnf_transformation, [], [f147])).
fof(f1123, plain, (~ spl78_79 | ~ spl78_16 | spl78_1), inference(avatar_split_clause, [], [f402, f690, f753, f1120])).
fof(f402, plain, ((e0 = op(e3, e3)) | ~ (e3 = op(e3, e0)) | ~ sP63), inference(cnf_transformation, [], [f147])).
fof(f1118, plain, (~ spl78_78 | spl78_57), inference(avatar_split_clause, [], [f391, f928, f1111])).
fof(f391, plain, ((e0 = op(e0, e1)) | ~ sP64), inference(cnf_transformation, [], [f146])).
fof(f146, plain, ((((e1 = op(e3, e3)) | ~ (e3 = op(e3, e1))) & ((e1 = op(e2, e2)) | ~ (e2 = op(e2, e1))) & ((e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & ((op(e0, e0) = e1) | ~ (e0 = op(e0, e1))) & (e0 = op(e1, e0)) & (e0 = op(e0, e1))) | ~ sP64), inference(nnf_transformation, [], [f118])).
fof(f1117, plain, (~ spl78_78 | spl78_45), inference(avatar_split_clause, [], [f392, f877, f1111])).
fof(f392, plain, ((e0 = op(e1, e0)) | ~ sP64), inference(cnf_transformation, [], [f146])).
fof(f1116, plain, (~ spl78_78 | ~ spl78_57 | spl78_62), inference(avatar_split_clause, [], [f393, f949, f928, f1111])).
fof(f393, plain, ((op(e0, e0) = e1) | ~ (e0 = op(e0, e1)) | ~ sP64), inference(cnf_transformation, [], [f146])).
fof(f1115, plain, (~ spl78_78 | ~ spl78_27 | spl78_22), inference(avatar_split_clause, [], [f395, f779, f800, f1111])).
fof(f395, plain, ((e1 = op(e2, e2)) | ~ (e2 = op(e2, e1)) | ~ sP64), inference(cnf_transformation, [], [f146])).
fof(f1114, plain, (~ spl78_78 | ~ spl78_12 | spl78_2), inference(avatar_split_clause, [], [f396, f694, f736, f1111])).
fof(f396, plain, ((e1 = op(e3, e3)) | ~ (e3 = op(e3, e1)) | ~ sP64), inference(cnf_transformation, [], [f146])).
fof(f1109, plain, (~ spl78_77 | spl78_53), inference(avatar_split_clause, [], [f385, f911, f1102])).
fof(f385, plain, ((e0 = op(e0, e2)) | ~ sP65), inference(cnf_transformation, [], [f145])).
fof(f145, plain, ((((e2 = op(e3, e3)) | ~ (e3 = op(e3, e2))) & ((e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & ((e2 = op(e1, e1)) | ~ (e1 = op(e1, e2))) & ((op(e0, e0) = e2) | ~ (e0 = op(e0, e2))) & (e0 = op(e2, e0)) & (e0 = op(e0, e2))) | ~ sP65), inference(nnf_transformation, [], [f119])).
fof(f1108, plain, (~ spl78_77 | spl78_29), inference(avatar_split_clause, [], [f386, f809, f1102])).
fof(f386, plain, ((e0 = op(e2, e0)) | ~ sP65), inference(cnf_transformation, [], [f145])).
fof(f1107, plain, (~ spl78_77 | ~ spl78_53 | spl78_63), inference(avatar_split_clause, [], [f387, f953, f911, f1102])).
fof(f387, plain, ((op(e0, e0) = e2) | ~ (e0 = op(e0, e2)) | ~ sP65), inference(cnf_transformation, [], [f145])).
fof(f1106, plain, (~ spl78_77 | ~ spl78_38 | spl78_43), inference(avatar_split_clause, [], [f388, f868, f847, f1102])).
fof(f388, plain, ((e2 = op(e1, e1)) | ~ (e1 = op(e1, e2)) | ~ sP65), inference(cnf_transformation, [], [f145])).
fof(f1105, plain, (~ spl78_77 | ~ spl78_8 | spl78_3), inference(avatar_split_clause, [], [f390, f698, f719, f1102])).
fof(f390, plain, ((e2 = op(e3, e3)) | ~ (e3 = op(e3, e2)) | ~ sP65), inference(cnf_transformation, [], [f145])).
fof(f1100, plain, (~ spl78_76 | spl78_49), inference(avatar_split_clause, [], [f379, f894, f1093])).
fof(f379, plain, ((e0 = op(e0, e3)) | ~ sP66), inference(cnf_transformation, [], [f144])).
fof(f144, plain, ((((e3 = op(e3, e3)) | ~ (e3 = op(e3, e3))) & ((e3 = op(e2, e2)) | ~ (e2 = op(e2, e3))) & ((e3 = op(e1, e1)) | ~ (e1 = op(e1, e3))) & ((op(e0, e0) = e3) | ~ (e0 = op(e0, e3))) & (e0 = op(e3, e0)) & (e0 = op(e0, e3))) | ~ sP66), inference(nnf_transformation, [], [f120])).
fof(f1099, plain, (~ spl78_76 | spl78_13), inference(avatar_split_clause, [], [f380, f741, f1093])).
fof(f380, plain, ((e0 = op(e3, e0)) | ~ sP66), inference(cnf_transformation, [], [f144])).
fof(f1098, plain, (~ spl78_76 | ~ spl78_49 | spl78_64), inference(avatar_split_clause, [], [f381, f957, f894, f1093])).
fof(f381, plain, ((op(e0, e0) = e3) | ~ (e0 = op(e0, e3)) | ~ sP66), inference(cnf_transformation, [], [f144])).
fof(f1097, plain, (~ spl78_76 | ~ spl78_34 | spl78_44), inference(avatar_split_clause, [], [f382, f872, f830, f1093])).
fof(f382, plain, ((e3 = op(e1, e1)) | ~ (e1 = op(e1, e3)) | ~ sP66), inference(cnf_transformation, [], [f144])).
fof(f1096, plain, (~ spl78_76 | ~ spl78_19 | spl78_24), inference(avatar_split_clause, [], [f383, f787, f766, f1093])).
fof(f383, plain, ((e3 = op(e2, e2)) | ~ (e2 = op(e2, e3)) | ~ sP66), inference(cnf_transformation, [], [f144])).
fof(f1091, plain, (~ spl78_75 | spl78_46), inference(avatar_split_clause, [], [f373, f881, f1084])).
fof(f373, plain, ((e1 = op(e1, e0)) | ~ sP67), inference(cnf_transformation, [], [f143])).
fof(f143, plain, ((((e0 = op(e3, e3)) | ~ (e3 = op(e3, e0))) & ((e0 = op(e2, e2)) | ~ (e2 = op(e2, e0))) & ((e0 = op(e1, e1)) | ~ (e1 = op(e1, e0))) & ((e0 = op(e0, e0)) | ~ (e0 = op(e0, e0))) & (e1 = op(e0, e1)) & (e1 = op(e1, e0))) | ~ sP67), inference(nnf_transformation, [], [f121])).
fof(f1090, plain, (~ spl78_75 | spl78_58), inference(avatar_split_clause, [], [f374, f932, f1084])).
fof(f374, plain, ((e1 = op(e0, e1)) | ~ sP67), inference(cnf_transformation, [], [f143])).
fof(f1089, plain, (~ spl78_75 | ~ spl78_46 | spl78_41), inference(avatar_split_clause, [], [f376, f860, f881, f1084])).
fof(f376, plain, ((e0 = op(e1, e1)) | ~ (e1 = op(e1, e0)) | ~ sP67), inference(cnf_transformation, [], [f143])).
fof(f1088, plain, (~ spl78_75 | ~ spl78_31 | spl78_21), inference(avatar_split_clause, [], [f377, f775, f817, f1084])).
fof(f377, plain, ((e0 = op(e2, e2)) | ~ (e2 = op(e2, e0)) | ~ sP67), inference(cnf_transformation, [], [f143])).
fof(f1087, plain, (~ spl78_75 | ~ spl78_16 | spl78_1), inference(avatar_split_clause, [], [f378, f690, f753, f1084])).
fof(f378, plain, ((e0 = op(e3, e3)) | ~ (e3 = op(e3, e0)) | ~ sP67), inference(cnf_transformation, [], [f143])).
fof(f1082, plain, (~ spl78_74 | spl78_42), inference(avatar_split_clause, [], [f367, f864, f1075])).
fof(f367, plain, ((e1 = op(e1, e1)) | ~ sP68), inference(cnf_transformation, [], [f142])).
fof(f142, plain, ((((e1 = op(e3, e3)) | ~ (e3 = op(e3, e1))) & ((e1 = op(e2, e2)) | ~ (e2 = op(e2, e1))) & ((e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & ((op(e0, e0) = e1) | ~ (e0 = op(e0, e1))) & (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP68), inference(nnf_transformation, [], [f122])).
fof(f1080, plain, (~ spl78_74 | ~ spl78_57 | spl78_62), inference(avatar_split_clause, [], [f369, f949, f928, f1075])).
fof(f369, plain, ((op(e0, e0) = e1) | ~ (e0 = op(e0, e1)) | ~ sP68), inference(cnf_transformation, [], [f142])).
fof(f1079, plain, (~ spl78_74 | ~ spl78_27 | spl78_22), inference(avatar_split_clause, [], [f371, f779, f800, f1075])).
fof(f371, plain, ((e1 = op(e2, e2)) | ~ (e2 = op(e2, e1)) | ~ sP68), inference(cnf_transformation, [], [f142])).
fof(f1078, plain, (~ spl78_74 | ~ spl78_12 | spl78_2), inference(avatar_split_clause, [], [f372, f694, f736, f1075])).
fof(f372, plain, ((e1 = op(e3, e3)) | ~ (e3 = op(e3, e1)) | ~ sP68), inference(cnf_transformation, [], [f142])).
fof(f1073, plain, (~ spl78_73 | spl78_38), inference(avatar_split_clause, [], [f361, f847, f1066])).
fof(f361, plain, ((e1 = op(e1, e2)) | ~ sP69), inference(cnf_transformation, [], [f141])).
fof(f141, plain, ((((e2 = op(e3, e3)) | ~ (e3 = op(e3, e2))) & ((e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & ((e2 = op(e1, e1)) | ~ (e1 = op(e1, e2))) & ((op(e0, e0) = e2) | ~ (e0 = op(e0, e2))) & (e1 = op(e2, e1)) & (e1 = op(e1, e2))) | ~ sP69), inference(nnf_transformation, [], [f123])).
fof(f1072, plain, (~ spl78_73 | spl78_26), inference(avatar_split_clause, [], [f362, f796, f1066])).
fof(f362, plain, ((e1 = op(e2, e1)) | ~ sP69), inference(cnf_transformation, [], [f141])).
fof(f1071, plain, (~ spl78_73 | ~ spl78_53 | spl78_63), inference(avatar_split_clause, [], [f363, f953, f911, f1066])).
fof(f363, plain, ((op(e0, e0) = e2) | ~ (e0 = op(e0, e2)) | ~ sP69), inference(cnf_transformation, [], [f141])).
fof(f1070, plain, (~ spl78_73 | ~ spl78_38 | spl78_43), inference(avatar_split_clause, [], [f364, f868, f847, f1066])).
fof(f364, plain, ((e2 = op(e1, e1)) | ~ (e1 = op(e1, e2)) | ~ sP69), inference(cnf_transformation, [], [f141])).
fof(f1069, plain, (~ spl78_73 | ~ spl78_8 | spl78_3), inference(avatar_split_clause, [], [f366, f698, f719, f1066])).
fof(f366, plain, ((e2 = op(e3, e3)) | ~ (e3 = op(e3, e2)) | ~ sP69), inference(cnf_transformation, [], [f141])).
fof(f1064, plain, (~ spl78_72 | spl78_34), inference(avatar_split_clause, [], [f355, f830, f1057])).
fof(f355, plain, ((e1 = op(e1, e3)) | ~ sP70), inference(cnf_transformation, [], [f140])).
fof(f140, plain, ((((e3 = op(e3, e3)) | ~ (e3 = op(e3, e3))) & ((e3 = op(e2, e2)) | ~ (e2 = op(e2, e3))) & ((e3 = op(e1, e1)) | ~ (e1 = op(e1, e3))) & ((op(e0, e0) = e3) | ~ (e0 = op(e0, e3))) & (e1 = op(e3, e1)) & (e1 = op(e1, e3))) | ~ sP70), inference(nnf_transformation, [], [f124])).
fof(f1063, plain, (~ spl78_72 | spl78_10), inference(avatar_split_clause, [], [f356, f728, f1057])).
fof(f356, plain, ((e1 = op(e3, e1)) | ~ sP70), inference(cnf_transformation, [], [f140])).
fof(f1062, plain, (~ spl78_72 | ~ spl78_49 | spl78_64), inference(avatar_split_clause, [], [f357, f957, f894, f1057])).
fof(f357, plain, ((op(e0, e0) = e3) | ~ (e0 = op(e0, e3)) | ~ sP70), inference(cnf_transformation, [], [f140])).
fof(f1061, plain, (~ spl78_72 | ~ spl78_34 | spl78_44), inference(avatar_split_clause, [], [f358, f872, f830, f1057])).
fof(f358, plain, ((e3 = op(e1, e1)) | ~ (e1 = op(e1, e3)) | ~ sP70), inference(cnf_transformation, [], [f140])).
fof(f1060, plain, (~ spl78_72 | ~ spl78_19 | spl78_24), inference(avatar_split_clause, [], [f359, f787, f766, f1057])).
fof(f359, plain, ((e3 = op(e2, e2)) | ~ (e2 = op(e2, e3)) | ~ sP70), inference(cnf_transformation, [], [f140])).
fof(f1055, plain, (~ spl78_71 | spl78_31), inference(avatar_split_clause, [], [f349, f817, f1048])).
fof(f349, plain, ((e2 = op(e2, e0)) | ~ sP71), inference(cnf_transformation, [], [f139])).
fof(f139, plain, ((((e0 = op(e3, e3)) | ~ (e3 = op(e3, e0))) & ((e0 = op(e2, e2)) | ~ (e2 = op(e2, e0))) & ((e0 = op(e1, e1)) | ~ (e1 = op(e1, e0))) & ((e0 = op(e0, e0)) | ~ (e0 = op(e0, e0))) & (e2 = op(e0, e2)) & (e2 = op(e2, e0))) | ~ sP71), inference(nnf_transformation, [], [f125])).
fof(f1054, plain, (~ spl78_71 | spl78_55), inference(avatar_split_clause, [], [f350, f919, f1048])).
fof(f350, plain, ((e2 = op(e0, e2)) | ~ sP71), inference(cnf_transformation, [], [f139])).
fof(f1053, plain, (~ spl78_71 | ~ spl78_46 | spl78_41), inference(avatar_split_clause, [], [f352, f860, f881, f1048])).
fof(f352, plain, ((e0 = op(e1, e1)) | ~ (e1 = op(e1, e0)) | ~ sP71), inference(cnf_transformation, [], [f139])).
fof(f1052, plain, (~ spl78_71 | ~ spl78_31 | spl78_21), inference(avatar_split_clause, [], [f353, f775, f817, f1048])).
fof(f353, plain, ((e0 = op(e2, e2)) | ~ (e2 = op(e2, e0)) | ~ sP71), inference(cnf_transformation, [], [f139])).
fof(f1051, plain, (~ spl78_71 | ~ spl78_16 | spl78_1), inference(avatar_split_clause, [], [f354, f690, f753, f1048])).
fof(f354, plain, ((e0 = op(e3, e3)) | ~ (e3 = op(e3, e0)) | ~ sP71), inference(cnf_transformation, [], [f139])).
fof(f1046, plain, (~ spl78_70 | spl78_27), inference(avatar_split_clause, [], [f343, f800, f1039])).
fof(f343, plain, ((e2 = op(e2, e1)) | ~ sP72), inference(cnf_transformation, [], [f138])).
fof(f138, plain, ((((e1 = op(e3, e3)) | ~ (e3 = op(e3, e1))) & ((e1 = op(e2, e2)) | ~ (e2 = op(e2, e1))) & ((e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & ((op(e0, e0) = e1) | ~ (e0 = op(e0, e1))) & (e2 = op(e1, e2)) & (e2 = op(e2, e1))) | ~ sP72), inference(nnf_transformation, [], [f126])).
fof(f1045, plain, (~ spl78_70 | spl78_39), inference(avatar_split_clause, [], [f344, f851, f1039])).
fof(f344, plain, ((e2 = op(e1, e2)) | ~ sP72), inference(cnf_transformation, [], [f138])).
fof(f1044, plain, (~ spl78_70 | ~ spl78_57 | spl78_62), inference(avatar_split_clause, [], [f345, f949, f928, f1039])).
fof(f345, plain, ((op(e0, e0) = e1) | ~ (e0 = op(e0, e1)) | ~ sP72), inference(cnf_transformation, [], [f138])).
fof(f1043, plain, (~ spl78_70 | ~ spl78_27 | spl78_22), inference(avatar_split_clause, [], [f347, f779, f800, f1039])).
fof(f347, plain, ((e1 = op(e2, e2)) | ~ (e2 = op(e2, e1)) | ~ sP72), inference(cnf_transformation, [], [f138])).
fof(f1042, plain, (~ spl78_70 | ~ spl78_12 | spl78_2), inference(avatar_split_clause, [], [f348, f694, f736, f1039])).
fof(f348, plain, ((e1 = op(e3, e3)) | ~ (e3 = op(e3, e1)) | ~ sP72), inference(cnf_transformation, [], [f138])).
fof(f1037, plain, (~ spl78_69 | spl78_23), inference(avatar_split_clause, [], [f337, f783, f1030])).
fof(f337, plain, ((e2 = op(e2, e2)) | ~ sP73), inference(cnf_transformation, [], [f137])).
fof(f137, plain, ((((e2 = op(e3, e3)) | ~ (e3 = op(e3, e2))) & ((e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & ((e2 = op(e1, e1)) | ~ (e1 = op(e1, e2))) & ((op(e0, e0) = e2) | ~ (e0 = op(e0, e2))) & (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP73), inference(nnf_transformation, [], [f127])).
fof(f1035, plain, (~ spl78_69 | ~ spl78_53 | spl78_63), inference(avatar_split_clause, [], [f339, f953, f911, f1030])).
fof(f339, plain, ((op(e0, e0) = e2) | ~ (e0 = op(e0, e2)) | ~ sP73), inference(cnf_transformation, [], [f137])).
fof(f1034, plain, (~ spl78_69 | ~ spl78_38 | spl78_43), inference(avatar_split_clause, [], [f340, f868, f847, f1030])).
fof(f340, plain, ((e2 = op(e1, e1)) | ~ (e1 = op(e1, e2)) | ~ sP73), inference(cnf_transformation, [], [f137])).
fof(f1033, plain, (~ spl78_69 | ~ spl78_8 | spl78_3), inference(avatar_split_clause, [], [f342, f698, f719, f1030])).
fof(f342, plain, ((e2 = op(e3, e3)) | ~ (e3 = op(e3, e2)) | ~ sP73), inference(cnf_transformation, [], [f137])).
fof(f1028, plain, (~ spl78_68 | spl78_19), inference(avatar_split_clause, [], [f331, f766, f1021])).
fof(f331, plain, ((e2 = op(e2, e3)) | ~ sP74), inference(cnf_transformation, [], [f136])).
fof(f136, plain, ((((e3 = op(e3, e3)) | ~ (e3 = op(e3, e3))) & ((e3 = op(e2, e2)) | ~ (e2 = op(e2, e3))) & ((e3 = op(e1, e1)) | ~ (e1 = op(e1, e3))) & ((op(e0, e0) = e3) | ~ (e0 = op(e0, e3))) & (e2 = op(e3, e2)) & (e2 = op(e2, e3))) | ~ sP74), inference(nnf_transformation, [], [f128])).
fof(f1027, plain, (~ spl78_68 | spl78_7), inference(avatar_split_clause, [], [f332, f715, f1021])).
fof(f332, plain, ((e2 = op(e3, e2)) | ~ sP74), inference(cnf_transformation, [], [f136])).
fof(f1026, plain, (~ spl78_68 | ~ spl78_49 | spl78_64), inference(avatar_split_clause, [], [f333, f957, f894, f1021])).
fof(f333, plain, ((op(e0, e0) = e3) | ~ (e0 = op(e0, e3)) | ~ sP74), inference(cnf_transformation, [], [f136])).
fof(f1025, plain, (~ spl78_68 | ~ spl78_34 | spl78_44), inference(avatar_split_clause, [], [f334, f872, f830, f1021])).
fof(f334, plain, ((e3 = op(e1, e1)) | ~ (e1 = op(e1, e3)) | ~ sP74), inference(cnf_transformation, [], [f136])).
fof(f1024, plain, (~ spl78_68 | ~ spl78_19 | spl78_24), inference(avatar_split_clause, [], [f335, f787, f766, f1021])).
fof(f335, plain, ((e3 = op(e2, e2)) | ~ (e2 = op(e2, e3)) | ~ sP74), inference(cnf_transformation, [], [f136])).
fof(f1019, plain, (~ spl78_67 | spl78_16), inference(avatar_split_clause, [], [f325, f753, f1012])).
fof(f325, plain, ((e3 = op(e3, e0)) | ~ sP75), inference(cnf_transformation, [], [f135])).
fof(f135, plain, ((((e0 = op(e3, e3)) | ~ (e3 = op(e3, e0))) & ((e0 = op(e2, e2)) | ~ (e2 = op(e2, e0))) & ((e0 = op(e1, e1)) | ~ (e1 = op(e1, e0))) & ((e0 = op(e0, e0)) | ~ (e0 = op(e0, e0))) & (e3 = op(e0, e3)) & (e3 = op(e3, e0))) | ~ sP75), inference(nnf_transformation, [], [f129])).
fof(f1018, plain, (~ spl78_67 | spl78_52), inference(avatar_split_clause, [], [f326, f906, f1012])).
fof(f326, plain, ((e3 = op(e0, e3)) | ~ sP75), inference(cnf_transformation, [], [f135])).
fof(f1017, plain, (~ spl78_67 | ~ spl78_46 | spl78_41), inference(avatar_split_clause, [], [f328, f860, f881, f1012])).
fof(f328, plain, ((e0 = op(e1, e1)) | ~ (e1 = op(e1, e0)) | ~ sP75), inference(cnf_transformation, [], [f135])).
fof(f1016, plain, (~ spl78_67 | ~ spl78_31 | spl78_21), inference(avatar_split_clause, [], [f329, f775, f817, f1012])).
fof(f329, plain, ((e0 = op(e2, e2)) | ~ (e2 = op(e2, e0)) | ~ sP75), inference(cnf_transformation, [], [f135])).
fof(f1015, plain, (~ spl78_67 | ~ spl78_16 | spl78_1), inference(avatar_split_clause, [], [f330, f690, f753, f1012])).
fof(f330, plain, ((e0 = op(e3, e3)) | ~ (e3 = op(e3, e0)) | ~ sP75), inference(cnf_transformation, [], [f135])).
fof(f1010, plain, (~ spl78_66 | spl78_12), inference(avatar_split_clause, [], [f319, f736, f1003])).
fof(f319, plain, ((e3 = op(e3, e1)) | ~ sP76), inference(cnf_transformation, [], [f134])).
fof(f134, plain, ((((e1 = op(e3, e3)) | ~ (e3 = op(e3, e1))) & ((e1 = op(e2, e2)) | ~ (e2 = op(e2, e1))) & ((e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & ((op(e0, e0) = e1) | ~ (e0 = op(e0, e1))) & (e3 = op(e1, e3)) & (e3 = op(e3, e1))) | ~ sP76), inference(nnf_transformation, [], [f130])).
fof(f1009, plain, (~ spl78_66 | spl78_36), inference(avatar_split_clause, [], [f320, f838, f1003])).
fof(f320, plain, ((e3 = op(e1, e3)) | ~ sP76), inference(cnf_transformation, [], [f134])).
fof(f1008, plain, (~ spl78_66 | ~ spl78_57 | spl78_62), inference(avatar_split_clause, [], [f321, f949, f928, f1003])).
fof(f321, plain, ((op(e0, e0) = e1) | ~ (e0 = op(e0, e1)) | ~ sP76), inference(cnf_transformation, [], [f134])).
fof(f1007, plain, (~ spl78_66 | ~ spl78_27 | spl78_22), inference(avatar_split_clause, [], [f323, f779, f800, f1003])).
fof(f323, plain, ((e1 = op(e2, e2)) | ~ (e2 = op(e2, e1)) | ~ sP76), inference(cnf_transformation, [], [f134])).
fof(f1006, plain, (~ spl78_66 | ~ spl78_12 | spl78_2), inference(avatar_split_clause, [], [f324, f694, f736, f1003])).
fof(f324, plain, ((e1 = op(e3, e3)) | ~ (e3 = op(e3, e1)) | ~ sP76), inference(cnf_transformation, [], [f134])).
fof(f1001, plain, (~ spl78_65 | spl78_8), inference(avatar_split_clause, [], [f313, f719, f994])).
fof(f313, plain, ((e3 = op(e3, e2)) | ~ sP77), inference(cnf_transformation, [], [f133])).
fof(f133, plain, ((((e2 = op(e3, e3)) | ~ (e3 = op(e3, e2))) & ((e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & ((e2 = op(e1, e1)) | ~ (e1 = op(e1, e2))) & ((op(e0, e0) = e2) | ~ (e0 = op(e0, e2))) & (e3 = op(e2, e3)) & (e3 = op(e3, e2))) | ~ sP77), inference(nnf_transformation, [], [f131])).
fof(f1000, plain, (~ spl78_65 | spl78_20), inference(avatar_split_clause, [], [f314, f770, f994])).
fof(f314, plain, ((e3 = op(e2, e3)) | ~ sP77), inference(cnf_transformation, [], [f133])).
fof(f999, plain, (~ spl78_65 | ~ spl78_53 | spl78_63), inference(avatar_split_clause, [], [f315, f953, f911, f994])).
fof(f315, plain, ((op(e0, e0) = e2) | ~ (e0 = op(e0, e2)) | ~ sP77), inference(cnf_transformation, [], [f133])).
fof(f998, plain, (~ spl78_65 | ~ spl78_38 | spl78_43), inference(avatar_split_clause, [], [f316, f868, f847, f994])).
fof(f316, plain, ((e2 = op(e1, e1)) | ~ (e1 = op(e1, e2)) | ~ sP77), inference(cnf_transformation, [], [f133])).
fof(f997, plain, (~ spl78_65 | ~ spl78_8 | spl78_3), inference(avatar_split_clause, [], [f318, f698, f719, f994])).
fof(f318, plain, ((e2 = op(e3, e3)) | ~ (e3 = op(e3, e2)) | ~ sP77), inference(cnf_transformation, [], [f133])).
fof(f992, plain, (spl78_61 | spl78_57 | spl78_53 | spl78_49), inference(avatar_split_clause, [], [f227, f894, f911, f928, f945])).
fof(f227, plain, ((e0 = op(e0, e3)) | (e0 = op(e0, e2)) | (e0 = op(e0, e1)) | (e0 = op(e0, e0))), inference(cnf_transformation, [], [f2])).
fof(f2, plain, (((e3 = op(e3, e3)) | (e3 = op(e2, e3)) | (e3 = op(e1, e3)) | (e3 = op(e0, e3))) & ((e3 = op(e3, e3)) | (e3 = op(e3, e2)) | (e3 = op(e3, e1)) | (e3 = op(e3, e0))) & ((e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(e0, e3))) & ((e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, e0))) & ((e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(e0, e3))) & ((e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, e0))) & ((e0 = op(e3, e3)) | (e0 = op(e2, e3)) | (e0 = op(e1, e3)) | (e0 = op(e0, e3))) & ((e0 = op(e3, e3)) | (e0 = op(e3, e2)) | (e0 = op(e3, e1)) | (e0 = op(e3, e0))) & ((e3 = op(e3, e2)) | (e3 = op(e2, e2)) | (e3 = op(e1, e2)) | (e3 = op(e0, e2))) & ((e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, e0))) & ((e2 = op(e3, e2)) | (e2 = op(e2, e2)) | (e2 = op(e1, e2)) | (e2 = op(e0, e2))) & ((e2 = op(e2, e3)) | (e2 = op(e2, e2)) | (e2 = op(e2, e1)) | (e2 = op(e2, e0))) & ((e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(e0, e2))) & ((e1 = op(e2, e3)) | (e1 = op(e2, e2)) | (e1 = op(e2, e1)) | (e1 = op(e2, e0))) & ((e0 = op(e3, e2)) | (e0 = op(e2, e2)) | (e0 = op(e1, e2)) | (e0 = op(e0, e2))) & ((e0 = op(e2, e3)) | (e0 = op(e2, e2)) | (e0 = op(e2, e1)) | (e0 = op(e2, e0))) & ((e3 = op(e3, e1)) | (e3 = op(e2, e1)) | (e3 = op(e1, e1)) | (e3 = op(e0, e1))) & ((e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, e0))) & ((e2 = op(e3, e1)) | (e2 = op(e2, e1)) | (e2 = op(e1, e1)) | (e2 = op(e0, e1))) & ((e2 = op(e1, e3)) | (e2 = op(e1, e2)) | (e2 = op(e1, e1)) | (e2 = op(e1, e0))) & ((e1 = op(e3, e1)) | (e1 = op(e2, e1)) | (e1 = op(e1, e1)) | (e1 = op(e0, e1))) & ((e1 = op(e1, e3)) | (e1 = op(e1, e2)) | (e1 = op(e1, e1)) | (e1 = op(e1, e0))) & ((e0 = op(e3, e1)) | (e0 = op(e2, e1)) | (e0 = op(e1, e1)) | (e0 = op(e0, e1))) & ((e0 = op(e1, e3)) | (e0 = op(e1, e2)) | (e0 = op(e1, e1)) | (e0 = op(e1, e0))) & ((e3 = op(e3, e0)) | (e3 = op(e2, e0)) | (e3 = op(e1, e0)) | (op(e0, e0) = e3)) & ((e3 = op(e0, e3)) | (e3 = op(e0, e2)) | (e3 = op(e0, e1)) | (op(e0, e0) = e3)) & ((e2 = op(e3, e0)) | (e2 = op(e2, e0)) | (e2 = op(e1, e0)) | (op(e0, e0) = e2)) & ((e2 = op(e0, e3)) | (e2 = op(e0, e2)) | (e2 = op(e0, e1)) | (op(e0, e0) = e2)) & ((e1 = op(e3, e0)) | (e1 = op(e2, e0)) | (e1 = op(e1, e0)) | (op(e0, e0) = e1)) & ((e1 = op(e0, e3)) | (e1 = op(e0, e2)) | (e1 = op(e0, e1)) | (op(e0, e0) = e1)) & ((e0 = op(e3, e0)) | (e0 = op(e2, e0)) | (e0 = op(e1, e0)) | (e0 = op(e0, e0))) & ((e0 = op(e0, e3)) | (e0 = op(e0, e2)) | (e0 = op(e0, e1)) | (e0 = op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG160+1.p', ax2)).
fof(f991, plain, (spl78_61 | spl78_45 | spl78_29 | spl78_13), inference(avatar_split_clause, [], [f228, f741, f809, f877, f945])).
fof(f228, plain, ((e0 = op(e3, e0)) | (e0 = op(e2, e0)) | (e0 = op(e1, e0)) | (e0 = op(e0, e0))), inference(cnf_transformation, [], [f2])).
fof(f990, plain, (spl78_62 | spl78_58 | spl78_54 | spl78_50), inference(avatar_split_clause, [], [f229, f898, f915, f932, f949])).
fof(f229, plain, ((e1 = op(e0, e3)) | (e1 = op(e0, e2)) | (e1 = op(e0, e1)) | (op(e0, e0) = e1)), inference(cnf_transformation, [], [f2])).
fof(f989, plain, (spl78_62 | spl78_46 | spl78_30 | spl78_14), inference(avatar_split_clause, [], [f230, f745, f813, f881, f949])).
fof(f230, plain, ((e1 = op(e3, e0)) | (e1 = op(e2, e0)) | (e1 = op(e1, e0)) | (op(e0, e0) = e1)), inference(cnf_transformation, [], [f2])).
fof(f988, plain, (spl78_63 | spl78_59 | spl78_55 | spl78_51), inference(avatar_split_clause, [], [f231, f902, f919, f936, f953])).
fof(f231, plain, ((e2 = op(e0, e3)) | (e2 = op(e0, e2)) | (e2 = op(e0, e1)) | (op(e0, e0) = e2)), inference(cnf_transformation, [], [f2])).
fof(f987, plain, (spl78_63 | spl78_47 | spl78_31 | spl78_15), inference(avatar_split_clause, [], [f232, f749, f817, f885, f953])).
fof(f232, plain, ((e2 = op(e3, e0)) | (e2 = op(e2, e0)) | (e2 = op(e1, e0)) | (op(e0, e0) = e2)), inference(cnf_transformation, [], [f2])).
fof(f986, plain, (spl78_64 | spl78_60 | spl78_56 | spl78_52), inference(avatar_split_clause, [], [f233, f906, f923, f940, f957])).
fof(f233, plain, ((e3 = op(e0, e3)) | (e3 = op(e0, e2)) | (e3 = op(e0, e1)) | (op(e0, e0) = e3)), inference(cnf_transformation, [], [f2])).
fof(f985, plain, (spl78_64 | spl78_48 | spl78_32 | spl78_16), inference(avatar_split_clause, [], [f234, f753, f821, f889, f957])).
fof(f234, plain, ((e3 = op(e3, e0)) | (e3 = op(e2, e0)) | (e3 = op(e1, e0)) | (op(e0, e0) = e3)), inference(cnf_transformation, [], [f2])).
fof(f984, plain, (spl78_45 | spl78_41 | spl78_37 | spl78_33), inference(avatar_split_clause, [], [f235, f826, f843, f860, f877])).
fof(f235, plain, ((e0 = op(e1, e3)) | (e0 = op(e1, e2)) | (e0 = op(e1, e1)) | (e0 = op(e1, e0))), inference(cnf_transformation, [], [f2])).
fof(f983, plain, (spl78_57 | spl78_41 | spl78_25 | spl78_9), inference(avatar_split_clause, [], [f236, f724, f792, f860, f928])).
fof(f236, plain, ((e0 = op(e3, e1)) | (e0 = op(e2, e1)) | (e0 = op(e1, e1)) | (e0 = op(e0, e1))), inference(cnf_transformation, [], [f2])).
fof(f982, plain, (spl78_46 | spl78_42 | spl78_38 | spl78_34), inference(avatar_split_clause, [], [f237, f830, f847, f864, f881])).
fof(f237, plain, ((e1 = op(e1, e3)) | (e1 = op(e1, e2)) | (e1 = op(e1, e1)) | (e1 = op(e1, e0))), inference(cnf_transformation, [], [f2])).
fof(f981, plain, (spl78_58 | spl78_42 | spl78_26 | spl78_10), inference(avatar_split_clause, [], [f238, f728, f796, f864, f932])).
fof(f238, plain, ((e1 = op(e3, e1)) | (e1 = op(e2, e1)) | (e1 = op(e1, e1)) | (e1 = op(e0, e1))), inference(cnf_transformation, [], [f2])).
fof(f980, plain, (spl78_47 | spl78_43 | spl78_39 | spl78_35), inference(avatar_split_clause, [], [f239, f834, f851, f868, f885])).
fof(f239, plain, ((e2 = op(e1, e3)) | (e2 = op(e1, e2)) | (e2 = op(e1, e1)) | (e2 = op(e1, e0))), inference(cnf_transformation, [], [f2])).
fof(f979, plain, (spl78_59 | spl78_43 | spl78_27 | spl78_11), inference(avatar_split_clause, [], [f240, f732, f800, f868, f936])).
fof(f240, plain, ((e2 = op(e3, e1)) | (e2 = op(e2, e1)) | (e2 = op(e1, e1)) | (e2 = op(e0, e1))), inference(cnf_transformation, [], [f2])).
fof(f978, plain, (spl78_48 | spl78_44 | spl78_40 | spl78_36), inference(avatar_split_clause, [], [f241, f838, f855, f872, f889])).
fof(f241, plain, ((e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, e0))), inference(cnf_transformation, [], [f2])).
fof(f977, plain, (spl78_60 | spl78_44 | spl78_28 | spl78_12), inference(avatar_split_clause, [], [f242, f736, f804, f872, f940])).
fof(f242, plain, ((e3 = op(e3, e1)) | (e3 = op(e2, e1)) | (e3 = op(e1, e1)) | (e3 = op(e0, e1))), inference(cnf_transformation, [], [f2])).
fof(f976, plain, (spl78_29 | spl78_25 | spl78_21 | spl78_17), inference(avatar_split_clause, [], [f243, f758, f775, f792, f809])).
fof(f243, plain, ((e0 = op(e2, e3)) | (e0 = op(e2, e2)) | (e0 = op(e2, e1)) | (e0 = op(e2, e0))), inference(cnf_transformation, [], [f2])).
fof(f975, plain, (spl78_53 | spl78_37 | spl78_21 | spl78_5), inference(avatar_split_clause, [], [f244, f707, f775, f843, f911])).
fof(f244, plain, ((e0 = op(e3, e2)) | (e0 = op(e2, e2)) | (e0 = op(e1, e2)) | (e0 = op(e0, e2))), inference(cnf_transformation, [], [f2])).
fof(f973, plain, (spl78_54 | spl78_38 | spl78_22 | spl78_6), inference(avatar_split_clause, [], [f246, f711, f779, f847, f915])).
fof(f246, plain, ((e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(e0, e2))), inference(cnf_transformation, [], [f2])).
fof(f972, plain, (spl78_31 | spl78_27 | spl78_23 | spl78_19), inference(avatar_split_clause, [], [f247, f766, f783, f800, f817])).
fof(f247, plain, ((e2 = op(e2, e3)) | (e2 = op(e2, e2)) | (e2 = op(e2, e1)) | (e2 = op(e2, e0))), inference(cnf_transformation, [], [f2])).
fof(f971, plain, (spl78_55 | spl78_39 | spl78_23 | spl78_7), inference(avatar_split_clause, [], [f248, f715, f783, f851, f919])).
fof(f248, plain, ((e2 = op(e3, e2)) | (e2 = op(e2, e2)) | (e2 = op(e1, e2)) | (e2 = op(e0, e2))), inference(cnf_transformation, [], [f2])).
fof(f970, plain, (spl78_32 | spl78_28 | spl78_24 | spl78_20), inference(avatar_split_clause, [], [f249, f770, f787, f804, f821])).
fof(f249, plain, ((e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, e0))), inference(cnf_transformation, [], [f2])).
fof(f969, plain, (spl78_56 | spl78_40 | spl78_24 | spl78_8), inference(avatar_split_clause, [], [f250, f719, f787, f855, f923])).
fof(f250, plain, ((e3 = op(e3, e2)) | (e3 = op(e2, e2)) | (e3 = op(e1, e2)) | (e3 = op(e0, e2))), inference(cnf_transformation, [], [f2])).
fof(f968, plain, (spl78_13 | spl78_9 | spl78_5 | spl78_1), inference(avatar_split_clause, [], [f251, f690, f707, f724, f741])).
fof(f251, plain, ((e0 = op(e3, e3)) | (e0 = op(e3, e2)) | (e0 = op(e3, e1)) | (e0 = op(e3, e0))), inference(cnf_transformation, [], [f2])).
fof(f967, plain, (spl78_49 | spl78_33 | spl78_17 | spl78_1), inference(avatar_split_clause, [], [f252, f690, f758, f826, f894])).
fof(f252, plain, ((e0 = op(e3, e3)) | (e0 = op(e2, e3)) | (e0 = op(e1, e3)) | (e0 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f966, plain, (spl78_14 | spl78_10 | spl78_6 | spl78_2), inference(avatar_split_clause, [], [f253, f694, f711, f728, f745])).
fof(f253, plain, ((e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, e0))), inference(cnf_transformation, [], [f2])).
fof(f965, plain, (spl78_50 | spl78_34 | spl78_18 | spl78_2), inference(avatar_split_clause, [], [f254, f694, f762, f830, f898])).
fof(f254, plain, ((e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f964, plain, (spl78_15 | spl78_11 | spl78_7 | spl78_3), inference(avatar_split_clause, [], [f255, f698, f715, f732, f749])).
fof(f255, plain, ((e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, e0))), inference(cnf_transformation, [], [f2])).
fof(f963, plain, (spl78_51 | spl78_35 | spl78_19 | spl78_3), inference(avatar_split_clause, [], [f256, f698, f766, f834, f902])).
fof(f256, plain, ((e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f962, plain, (spl78_16 | spl78_12 | spl78_8 | spl78_4), inference(avatar_split_clause, [], [f257, f702, f719, f736, f753])).
fof(f257, plain, ((e3 = op(e3, e3)) | (e3 = op(e3, e2)) | (e3 = op(e3, e1)) | (e3 = op(e3, e0))), inference(cnf_transformation, [], [f2])).
fof(f961, plain, (spl78_52 | spl78_36 | spl78_20 | spl78_4), inference(avatar_split_clause, [], [f258, f702, f770, f838, f906])).
fof(f258, plain, ((e3 = op(e3, e3)) | (e3 = op(e2, e3)) | (e3 = op(e1, e3)) | (e3 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f960, plain, (spl78_61 | spl78_62 | spl78_63 | spl78_64), inference(avatar_split_clause, [], [f211, f957, f953, f949, f945])).
fof(f211, plain, ((op(e0, e0) = e3) | (op(e0, e0) = e2) | (op(e0, e0) = e1) | (e0 = op(e0, e0))), inference(cnf_transformation, [], [f1])).
fof(f1, plain, (((e3 = op(e3, e3)) | (e2 = op(e3, e3)) | (e1 = op(e3, e3)) | (e0 = op(e3, e3))) & ((e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (e0 = op(e3, e2))) & ((e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (e0 = op(e3, e1))) & ((e3 = op(e3, e0)) | (e2 = op(e3, e0)) | (e1 = op(e3, e0)) | (e0 = op(e3, e0))) & ((e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))) & ((e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (e0 = op(e2, e2))) & ((e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))) & ((e3 = op(e2, e0)) | (e2 = op(e2, e0)) | (e1 = op(e2, e0)) | (e0 = op(e2, e0))) & ((e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (e0 = op(e1, e3))) & ((e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (e0 = op(e1, e2))) & ((e3 = op(e1, e1)) | (e2 = op(e1, e1)) | (e1 = op(e1, e1)) | (e0 = op(e1, e1))) & ((e3 = op(e1, e0)) | (e2 = op(e1, e0)) | (e1 = op(e1, e0)) | (e0 = op(e1, e0))) & ((e3 = op(e0, e3)) | (e2 = op(e0, e3)) | (e1 = op(e0, e3)) | (e0 = op(e0, e3))) & ((e3 = op(e0, e2)) | (e2 = op(e0, e2)) | (e1 = op(e0, e2)) | (e0 = op(e0, e2))) & ((e3 = op(e0, e1)) | (e2 = op(e0, e1)) | (e1 = op(e0, e1)) | (e0 = op(e0, e1))) & ((op(e0, e0) = e3) | (op(e0, e0) = e2) | (op(e0, e0) = e1) | (e0 = op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG160+1.p', ax1)).
fof(f943, plain, (spl78_57 | spl78_58 | spl78_59 | spl78_60), inference(avatar_split_clause, [], [f212, f940, f936, f932, f928])).
fof(f212, plain, ((e3 = op(e0, e1)) | (e2 = op(e0, e1)) | (e1 = op(e0, e1)) | (e0 = op(e0, e1))), inference(cnf_transformation, [], [f1])).
fof(f926, plain, (spl78_53 | spl78_54 | spl78_55 | spl78_56), inference(avatar_split_clause, [], [f213, f923, f919, f915, f911])).
fof(f213, plain, ((e3 = op(e0, e2)) | (e2 = op(e0, e2)) | (e1 = op(e0, e2)) | (e0 = op(e0, e2))), inference(cnf_transformation, [], [f1])).
fof(f909, plain, (spl78_49 | spl78_50 | spl78_51 | spl78_52), inference(avatar_split_clause, [], [f214, f906, f902, f898, f894])).
fof(f214, plain, ((e3 = op(e0, e3)) | (e2 = op(e0, e3)) | (e1 = op(e0, e3)) | (e0 = op(e0, e3))), inference(cnf_transformation, [], [f1])).
fof(f892, plain, (spl78_45 | spl78_46 | spl78_47 | spl78_48), inference(avatar_split_clause, [], [f215, f889, f885, f881, f877])).
fof(f215, plain, ((e3 = op(e1, e0)) | (e2 = op(e1, e0)) | (e1 = op(e1, e0)) | (e0 = op(e1, e0))), inference(cnf_transformation, [], [f1])).
fof(f875, plain, (spl78_41 | spl78_42 | spl78_43 | spl78_44), inference(avatar_split_clause, [], [f216, f872, f868, f864, f860])).
fof(f216, plain, ((e3 = op(e1, e1)) | (e2 = op(e1, e1)) | (e1 = op(e1, e1)) | (e0 = op(e1, e1))), inference(cnf_transformation, [], [f1])).
fof(f858, plain, (spl78_37 | spl78_38 | spl78_39 | spl78_40), inference(avatar_split_clause, [], [f217, f855, f851, f847, f843])).
fof(f217, plain, ((e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (e0 = op(e1, e2))), inference(cnf_transformation, [], [f1])).
fof(f841, plain, (spl78_33 | spl78_34 | spl78_35 | spl78_36), inference(avatar_split_clause, [], [f218, f838, f834, f830, f826])).
fof(f218, plain, ((e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (e0 = op(e1, e3))), inference(cnf_transformation, [], [f1])).
fof(f824, plain, (spl78_29 | spl78_30 | spl78_31 | spl78_32), inference(avatar_split_clause, [], [f219, f821, f817, f813, f809])).
fof(f219, plain, ((e3 = op(e2, e0)) | (e2 = op(e2, e0)) | (e1 = op(e2, e0)) | (e0 = op(e2, e0))), inference(cnf_transformation, [], [f1])).
fof(f807, plain, (spl78_25 | spl78_26 | spl78_27 | spl78_28), inference(avatar_split_clause, [], [f220, f804, f800, f796, f792])).
fof(f220, plain, ((e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))), inference(cnf_transformation, [], [f1])).
fof(f790, plain, (spl78_21 | spl78_22 | spl78_23 | spl78_24), inference(avatar_split_clause, [], [f221, f787, f783, f779, f775])).
fof(f221, plain, ((e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (e0 = op(e2, e2))), inference(cnf_transformation, [], [f1])).
fof(f773, plain, (spl78_17 | spl78_18 | spl78_19 | spl78_20), inference(avatar_split_clause, [], [f222, f770, f766, f762, f758])).
fof(f222, plain, ((e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))), inference(cnf_transformation, [], [f1])).
fof(f756, plain, (spl78_13 | spl78_14 | spl78_15 | spl78_16), inference(avatar_split_clause, [], [f223, f753, f749, f745, f741])).
fof(f223, plain, ((e3 = op(e3, e0)) | (e2 = op(e3, e0)) | (e1 = op(e3, e0)) | (e0 = op(e3, e0))), inference(cnf_transformation, [], [f1])).
fof(f739, plain, (spl78_9 | spl78_10 | spl78_11 | spl78_12), inference(avatar_split_clause, [], [f224, f736, f732, f728, f724])).
fof(f224, plain, ((e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (e0 = op(e3, e1))), inference(cnf_transformation, [], [f1])).
fof(f722, plain, (spl78_5 | spl78_6 | spl78_7 | spl78_8), inference(avatar_split_clause, [], [f225, f719, f715, f711, f707])).
fof(f225, plain, ((e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (e0 = op(e3, e2))), inference(cnf_transformation, [], [f1])).
fof(f705, plain, (spl78_1 | spl78_2 | spl78_3 | spl78_4), inference(avatar_split_clause, [], [f226, f702, f698, f694, f690])).
fof(f226, plain, ((e3 = op(e3, e3)) | (e2 = op(e3, e3)) | (e1 = op(e3, e3)) | (e0 = op(e3, e3))), inference(cnf_transformation, [], [f1])).