fof(f4149, plain, $false, inference(avatar_sat_refutation, [], [f605, f622, f639, f656, f673, f690, f707, f724, f741, f758, f775, f792, f809, f826, f843, f860, f861, f863, f864, f865, f866, f868, f870, f873, f875, f876, f878, f879, f880, f882, f883, f884, f885, f887, f888, f900, f906, f908, f916, f921, f923, f924, f930, f931, f937, f938, f939, f940, f945, f946, f947, f948, f953, f954, f955, f956, f961, f962, f963, f964, f970, f971, f972, f977, f978, f979, f980, f985, f986, f987, f988, f993, f994, f995, f996, f1001, f1002, f1003, f1004, f1010, f1018, f1025, f1026, f1028, f1033, f1034, f1035, f1036, f1041, f1042, f1043, f1044, f1049, f1052, f1060, f1065, f1068, f1073, f1074, f1076, f1081, f1082, f1083, f1084, f1089, f1090, f1091, f1092, f1098, f1099, f1100, f1105, f1106, f1107, f1108, f1113, f1114, f1115, f1116, f1121, f1122, f1123, f1124, f1129, f1130, f1131, f1132, f1138, f1146, f1148, f1153, f1154, f1155, f1156, f1161, f1162, f1163, f1164, f1169, f1170, f1171, f1172, f1177, f1178, f1179, f1180, f1186, f1188, f1193, f1194, f1195, f1196, f1201, f1202, f1203, f1204, f1209, f1210, f1212, f1217, f1220, f1228, f1233, f1236, f1241, f1242, f1243, f1244, f1249, f1250, f1251, f1252, f1257, f1258, f1260, f1266, f1274, f1276, f1281, f1282, f1283, f1284, f1289, f1290, f1291, f1292, f1297, f1298, f1299, f1300, f1305, f1306, f1307, f1308, f1314, f1316, f1321, f1322, f1323, f1324, f1329, f1330, f1331, f1332, f1337, f1338, f1339, f1340, f1345, f1346, f1347, f1348, f1354, f1356, f1361, f1363, f1364, f1372, f1380, f1385, f1388, f1396, f1405, f1410, f1415, f1420, f1434, f1439, f1444, f1458, f1463, f1468, f1472, f1477, f1492, f1493, f1494, f1495, f1496, f1505, f1514, f1573, f1578, f1583, f1588, f1593, f1608, f1609, f1610, f1611, f1614, f1615, f1616, f1628, f1641, f1642, f1643, f1644, f1649, f1650, f1661, f1664, f1670, f1671, f1674, f1699, f1702, f1705, f1734, f1737, f1747, f1751, f1763, f1773, f1776, f1779, f1784, f1797, f1806, f1819, f1823, f1826, f1834, f1840, f1848, f1849, f1855, f1884, f1898, f1904, f1926, f1937, f1950, f1951, f1959, f1966, f1986, f1990, f1999, f2012, f2038, f2060, f2067, f2070, f2080, f2096, f2116, f2118, f2128, f2156, f2173, f2177, f2182, f2193, f2198, f2205, f2225, f2238, f2241, f2290, f2311, f2316, f2337, f2349, f2364, f2390, f2393, f2405, f2420, f2429, f2453, f2475, f2504, f2518, f2519, f2594, f2595, f2596, f2620, f2624, f2639, f2644, f2655, f2686, f2693, f2697, f2698, f2704, f2709, f2729, f2738, f2748, f2754, f2762, f2775, f2782, f2812, f2814, f2824, f2829, f2835, f2841, f2842, f2864, f2883, f2890, f2903, f2915, f2923, f2931, f2947, f2948, f2949, f2977, f2999, f3004, f3043, f3045, f3059, f3088, f3137, f3186, f3187, f3193, f3200, f3213, f3217, f3223, f3246, f3272, f3287, f3308, f3309, f3328, f3335, f3345, f3351, f3358, f3360, f3370, f3371, f3374, f3383, f3400, f3401, f3410, f3412, f3427, f3428, f3437, f3444, f3480, f3481, f3501, f3506, f3512, f3513, f3525, f3552, f3553, f3570, f3595, f3596, f3603, f3608, f3619, f3631, f3633, f3640, f3649, f3651, f3668, f3690, f3693, f3709, f3710, f3717, f3724, f3725, f3736, f3749, f3765, f3772, f3789, f3794, f3796, f3823, f3839, f3853, f3863, f3873, f3878, f3879, f3880, f3885, f3894, f3905, f3909, f3910, f3916, f3927, f3929, f3932, f3934, f3940, f3941, f3952, f3976, f3983, f3991, f3997, f4003, f4010, f4016, f4017, f4026, f4030, f4037, f4053, f4062, f4086, f4088, f4093, f4100, f4105, f4112, f4113, f4117, f4134, f4139, f4148])).
fof(f4148, plain, (~ spl66_2 | ~ spl66_59 | spl66_157 | ~ spl66_166), inference(avatar_contradiction_clause, [], [f4147])).
fof(f4147, plain, ($false | (~ spl66_2 | ~ spl66_59 | spl66_157 | ~ spl66_166)), inference(subsumption_resolution, [], [f4145, f838])).
fof(f838, plain, ((e2 = op(e0, e1)) | ~ spl66_59), inference(avatar_component_clause, [], [f836])).
fof(f836, plain, (spl66_59 <=> (e2 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_59])])).
fof(f4145, plain, (~ (e2 = op(e0, e1)) | (~ spl66_2 | spl66_157 | ~ spl66_166)), inference(backward_demodulation, [], [f4073, f596])).
fof(f596, plain, ((e1 = op(e3, e3)) | ~ spl66_2), inference(avatar_component_clause, [], [f594])).
fof(f594, plain, (spl66_2 <=> (e1 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_2])])).
fof(f4073, plain, (~ (e2 = op(e0, op(e3, e3))) | (spl66_157 | ~ spl66_166)), inference(backward_demodulation, [], [f1545, f1586])).
fof(f1586, plain, ((e0 = op(op(e3, e3), e3)) | ~ spl66_166), inference(avatar_component_clause, [], [f1585])).
fof(f1585, plain, (spl66_166 <=> (e0 = op(op(e3, e3), e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_166])])).
fof(f1545, plain, (~ (e2 = op(op(op(e3, e3), e3), op(e3, e3))) | spl66_157), inference(avatar_component_clause, [], [f1543])).
fof(f1543, plain, (spl66_157 <=> (e2 = op(op(op(e3, e3), e3), op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl66_157])])).
fof(f4139, plain, (~ spl66_7 | ~ spl66_11), inference(avatar_split_clause, [], [f4136, f632, f615])).
fof(f615, plain, (spl66_7 <=> (e2 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_7])])).
fof(f632, plain, (spl66_11 <=> (e2 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_11])])).
fof(f4136, plain, (~ (e2 = op(e3, e2)) | ~ spl66_11), inference(backward_demodulation, [], [f280, f634])).
fof(f634, plain, ((e2 = op(e3, e1)) | ~ spl66_11), inference(avatar_component_clause, [], [f632])).
fof(f280, plain, ~ (op(e3, e1) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f3, plain, (~ (op(e3, e2) = op(e3, e3)) & ~ (op(e3, e1) = op(e3, e3)) & ~ (op(e3, e1) = op(e3, e2)) & ~ (op(e3, e0) = op(e3, e3)) & ~ (op(e3, e0) = op(e3, e2)) & ~ (op(e3, e0) = op(e3, e1)) & ~ (op(e2, e2) = op(e2, e3)) & ~ (op(e2, e1) = op(e2, e3)) & ~ (op(e2, e1) = op(e2, e2)) & ~ (op(e2, e0) = op(e2, e3)) & ~ (op(e2, e0) = op(e2, e2)) & ~ (op(e2, e0) = op(e2, e1)) & ~ (op(e1, e2) = op(e1, e3)) & ~ (op(e1, e1) = op(e1, e3)) & ~ (op(e1, e1) = op(e1, e2)) & ~ (op(e1, e0) = op(e1, e3)) & ~ (op(e1, e0) = op(e1, e2)) & ~ (op(e1, e0) = op(e1, e1)) & ~ (op(e0, e2) = op(e0, e3)) & ~ (op(e0, e1) = op(e0, e3)) & ~ (op(e0, e1) = op(e0, e2)) & ~ (op(e0, e0) = op(e0, e3)) & ~ (op(e0, e0) = op(e0, e2)) & ~ (op(e0, e0) = op(e0, e1)) & ~ (op(e2, e3) = op(e3, e3)) & ~ (op(e1, e3) = op(e3, e3)) & ~ (op(e1, e3) = op(e2, e3)) & ~ (op(e0, e3) = op(e3, e3)) & ~ (op(e0, e3) = op(e2, e3)) & ~ (op(e0, e3) = op(e1, e3)) & ~ (op(e2, e2) = op(e3, e2)) & ~ (op(e1, e2) = op(e3, e2)) & ~ (op(e1, e2) = op(e2, e2)) & ~ (op(e0, e2) = op(e3, e2)) & ~ (op(e0, e2) = op(e2, e2)) & ~ (op(e0, e2) = op(e1, e2)) & ~ (op(e2, e1) = op(e3, e1)) & ~ (op(e1, e1) = op(e3, e1)) & ~ (op(e1, e1) = op(e2, e1)) & ~ (op(e0, e1) = op(e3, e1)) & ~ (op(e0, e1) = op(e2, e1)) & ~ (op(e0, e1) = op(e1, e1)) & ~ (op(e2, e0) = op(e3, e0)) & ~ (op(e1, e0) = op(e3, e0)) & ~ (op(e1, e0) = op(e2, e0)) & ~ (op(e0, e0) = op(e3, e0)) & ~ (op(e0, e0) = op(e2, e0)) & ~ (op(e0, e0) = op(e1, e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG143+1.p', ax3)).
fof(f4134, plain, (~ spl66_9 | ~ spl66_13), inference(avatar_split_clause, [], [f4127, f641, f624])).
fof(f624, plain, (spl66_9 <=> (e0 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_9])])).
fof(f641, plain, (spl66_13 <=> (e0 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_13])])).
fof(f4127, plain, (~ (e0 = op(e3, e1)) | ~ spl66_13), inference(backward_demodulation, [], [f277, f643])).
fof(f643, plain, ((e0 = op(e3, e0)) | ~ spl66_13), inference(avatar_component_clause, [], [f641])).
fof(f277, plain, ~ (op(e3, e0) = op(e3, e1)), inference(cnf_transformation, [], [f3])).
fof(f4117, plain, (~ spl66_20 | ~ spl66_28), inference(avatar_split_clause, [], [f4116, f704, f670])).
fof(f670, plain, (spl66_20 <=> (e3 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_20])])).
fof(f704, plain, (spl66_28 <=> (e3 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_28])])).
fof(f4116, plain, (~ (e3 = op(e2, e3)) | ~ spl66_28), inference(backward_demodulation, [], [f275, f706])).
fof(f706, plain, ((e3 = op(e2, e1)) | ~ spl66_28), inference(avatar_component_clause, [], [f704])).
fof(f275, plain, ~ (op(e2, e1) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f4113, plain, (~ spl66_19 | ~ spl66_31), inference(avatar_split_clause, [], [f4109, f717, f666])).
fof(f666, plain, (spl66_19 <=> (e2 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_19])])).
fof(f717, plain, (spl66_31 <=> (e2 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_31])])).
fof(f4109, plain, (~ (e2 = op(e2, e3)) | ~ spl66_31), inference(backward_demodulation, [], [f273, f719])).
fof(f719, plain, ((e2 = op(e2, e0)) | ~ spl66_31), inference(avatar_component_clause, [], [f717])).
fof(f273, plain, ~ (op(e2, e0) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f4112, plain, (~ spl66_27 | ~ spl66_31), inference(avatar_split_clause, [], [f4108, f717, f700])).
fof(f700, plain, (spl66_27 <=> (e2 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_27])])).
fof(f4108, plain, (~ (e2 = op(e2, e1)) | ~ spl66_31), inference(backward_demodulation, [], [f271, f719])).
fof(f271, plain, ~ (op(e2, e0) = op(e2, e1)), inference(cnf_transformation, [], [f3])).
fof(f4105, plain, (~ spl66_7 | ~ spl66_39), inference(avatar_split_clause, [], [f4102, f751, f615])).
fof(f751, plain, (spl66_39 <=> (e2 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_39])])).
fof(f4102, plain, (~ (e2 = op(e3, e2)) | ~ spl66_39), inference(backward_demodulation, [], [f251, f753])).
fof(f753, plain, ((e2 = op(e1, e2)) | ~ spl66_39), inference(avatar_component_clause, [], [f751])).
fof(f251, plain, ~ (op(e1, e2) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f4100, plain, (~ spl66_16 | ~ spl66_48), inference(avatar_split_clause, [], [f4095, f789, f653])).
fof(f653, plain, (spl66_16 <=> (e3 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_16])])).
fof(f789, plain, (spl66_48 <=> (e3 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_48])])).
fof(f4095, plain, (~ (e3 = op(e3, e0)) | ~ spl66_48), inference(backward_demodulation, [], [f239, f791])).
fof(f791, plain, ((e3 = op(e1, e0)) | ~ spl66_48), inference(avatar_component_clause, [], [f789])).
fof(f239, plain, ~ (op(e1, e0) = op(e3, e0)), inference(cnf_transformation, [], [f3])).
fof(f4093, plain, (~ spl66_20 | ~ spl66_52), inference(avatar_split_clause, [], [f4092, f806, f670])).
fof(f806, plain, (spl66_52 <=> (e3 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_52])])).
fof(f4092, plain, (~ (e3 = op(e2, e3)) | ~ spl66_52), inference(backward_demodulation, [], [f254, f808])).
fof(f808, plain, ((e3 = op(e0, e3)) | ~ spl66_52), inference(avatar_component_clause, [], [f806])).
fof(f254, plain, ~ (op(e0, e3) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f4088, plain, (~ spl66_51 | ~ spl66_59), inference(avatar_split_clause, [], [f4083, f836, f802])).
fof(f802, plain, (spl66_51 <=> (e2 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_51])])).
fof(f4083, plain, (~ (e2 = op(e0, e3)) | ~ spl66_59), inference(backward_demodulation, [], [f263, f838])).
fof(f263, plain, ~ (op(e0, e1) = op(e0, e3)), inference(cnf_transformation, [], [f3])).
fof(f4086, plain, (~ spl66_27 | ~ spl66_59), inference(avatar_split_clause, [], [f4081, f836, f700])).
fof(f4081, plain, (~ (e2 = op(e2, e1)) | ~ spl66_59), inference(backward_demodulation, [], [f242, f838])).
fof(f242, plain, ~ (op(e0, e1) = op(e2, e1)), inference(cnf_transformation, [], [f3])).
fof(f4062, plain, (~ spl66_1 | ~ spl66_33), inference(avatar_split_clause, [], [f4061, f726, f590])).
fof(f590, plain, (spl66_1 <=> (e0 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_1])])).
fof(f726, plain, (spl66_33 <=> (e0 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_33])])).
fof(f4061, plain, (~ (e0 = op(e3, e3)) | ~ spl66_33), inference(forward_demodulation, [], [f257, f728])).
fof(f728, plain, ((e0 = op(e1, e3)) | ~ spl66_33), inference(avatar_component_clause, [], [f726])).
fof(f257, plain, ~ (op(e1, e3) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f4053, plain, (~ spl66_14 | ~ spl66_62), inference(avatar_split_clause, [], [f4052, f849, f645])).
fof(f645, plain, (spl66_14 <=> (e1 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_14])])).
fof(f849, plain, (spl66_62 <=> (op(e0, e0) = e1)), introduced(avatar_definition, [new_symbols(naming, [spl66_62])])).
fof(f4052, plain, (~ (e1 = op(e3, e0)) | ~ spl66_62), inference(forward_demodulation, [], [f237, f851])).
fof(f851, plain, ((op(e0, e0) = e1) | ~ spl66_62), inference(avatar_component_clause, [], [f849])).
fof(f237, plain, ~ (op(e0, e0) = op(e3, e0)), inference(cnf_transformation, [], [f3])).
fof(f4037, plain, (~ spl66_1 | ~ spl66_9), inference(avatar_split_clause, [], [f4032, f624, f590])).
fof(f4032, plain, (~ (e0 = op(e3, e3)) | ~ spl66_9), inference(backward_demodulation, [], [f281, f626])).
fof(f626, plain, ((e0 = op(e3, e1)) | ~ spl66_9), inference(avatar_component_clause, [], [f624])).
fof(f281, plain, ~ (op(e3, e1) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f4030, plain, (~ spl66_14 | ~ spl66_16), inference(avatar_contradiction_clause, [], [f4029])).
fof(f4029, plain, ($false | (~ spl66_14 | ~ spl66_16)), inference(subsumption_resolution, [], [f4028, f287])).
fof(f287, plain, ~ (e1 = e3), inference(cnf_transformation, [], [f4])).
fof(f4, plain, (~ (e2 = e3) & ~ (e1 = e3) & ~ (e1 = e2) & ~ (e0 = e3) & ~ (e0 = e2) & ~ (e0 = e1)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG143+1.p', ax4)).
fof(f4028, plain, ((e1 = e3) | (~ spl66_14 | ~ spl66_16)), inference(backward_demodulation, [], [f655, f647])).
fof(f647, plain, ((e1 = op(e3, e0)) | ~ spl66_14), inference(avatar_component_clause, [], [f645])).
fof(f655, plain, ((e3 = op(e3, e0)) | ~ spl66_16), inference(avatar_component_clause, [], [f653])).
fof(f4026, plain, (~ spl66_9 | ~ spl66_22 | ~ spl66_152 | spl66_163), inference(avatar_split_clause, [], [f4024, f1570, f1520, f679, f624])).
fof(f679, plain, (spl66_22 <=> (e1 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_22])])).
fof(f1520, plain, (spl66_152 <=> (e3 = op(op(e2, e2), e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_152])])).
fof(f1570, plain, (spl66_163 <=> (e0 = op(op(op(e2, e2), e2), op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl66_163])])).
fof(f4024, plain, (~ (e0 = op(e3, e1)) | (~ spl66_22 | ~ spl66_152 | spl66_163)), inference(backward_demodulation, [], [f3960, f681])).
fof(f681, plain, ((e1 = op(e2, e2)) | ~ spl66_22), inference(avatar_component_clause, [], [f679])).
fof(f3960, plain, (~ (e0 = op(e3, op(e2, e2))) | (~ spl66_152 | spl66_163)), inference(backward_demodulation, [], [f1572, f1521])).
fof(f1521, plain, ((e3 = op(op(e2, e2), e2)) | ~ spl66_152), inference(avatar_component_clause, [], [f1520])).
fof(f1572, plain, (~ (e0 = op(op(op(e2, e2), e2), op(e2, e2))) | spl66_163), inference(avatar_component_clause, [], [f1570])).
fof(f4017, plain, (~ spl66_23 | ~ spl66_27), inference(avatar_split_clause, [], [f4013, f700, f683])).
fof(f683, plain, (spl66_23 <=> (e2 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_23])])).
fof(f4013, plain, (~ (e2 = op(e2, e2)) | ~ spl66_27), inference(backward_demodulation, [], [f274, f702])).
fof(f702, plain, ((e2 = op(e2, e1)) | ~ spl66_27), inference(avatar_component_clause, [], [f700])).
fof(f274, plain, ~ (op(e2, e1) = op(e2, e2)), inference(cnf_transformation, [], [f3])).
fof(f4016, plain, (~ spl66_11 | ~ spl66_27), inference(avatar_split_clause, [], [f4012, f700, f632])).
fof(f4012, plain, (~ (e2 = op(e3, e1)) | ~ spl66_27), inference(backward_demodulation, [], [f246, f702])).
fof(f246, plain, ~ (op(e2, e1) = op(e3, e1)), inference(cnf_transformation, [], [f3])).
fof(f4010, plain, (~ spl66_25 | ~ spl66_29), inference(avatar_split_clause, [], [f4007, f709, f692])).
fof(f692, plain, (spl66_25 <=> (e0 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_25])])).
fof(f709, plain, (spl66_29 <=> (e0 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_29])])).
fof(f4007, plain, (~ (e0 = op(e2, e1)) | ~ spl66_29), inference(backward_demodulation, [], [f271, f711])).
fof(f711, plain, ((e0 = op(e2, e0)) | ~ spl66_29), inference(avatar_component_clause, [], [f709])).
fof(f4003, plain, (~ spl66_33 | ~ spl66_34), inference(avatar_contradiction_clause, [], [f4002])).
fof(f4002, plain, ($false | (~ spl66_33 | ~ spl66_34)), inference(subsumption_resolution, [], [f4001, f283])).
fof(f283, plain, ~ (e0 = e1), inference(cnf_transformation, [], [f4])).
fof(f4001, plain, ((e0 = e1) | (~ spl66_33 | ~ spl66_34)), inference(backward_demodulation, [], [f732, f728])).
fof(f732, plain, ((e1 = op(e1, e3)) | ~ spl66_34), inference(avatar_component_clause, [], [f730])).
fof(f730, plain, (spl66_34 <=> (e1 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_34])])).
fof(f3997, plain, (~ spl66_2 | ~ spl66_34), inference(avatar_split_clause, [], [f3994, f730, f594])).
fof(f3994, plain, (~ (e1 = op(e3, e3)) | ~ spl66_34), inference(backward_demodulation, [], [f257, f732])).
fof(f3991, plain, (~ spl66_35 | ~ spl66_47), inference(avatar_split_clause, [], [f3985, f785, f734])).
fof(f734, plain, (spl66_35 <=> (e2 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_35])])).
fof(f785, plain, (spl66_47 <=> (e2 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_47])])).
fof(f3985, plain, (~ (e2 = op(e1, e3)) | ~ spl66_47), inference(backward_demodulation, [], [f267, f787])).
fof(f787, plain, ((e2 = op(e1, e0)) | ~ spl66_47), inference(avatar_component_clause, [], [f785])).
fof(f267, plain, ~ (op(e1, e0) = op(e1, e3)), inference(cnf_transformation, [], [f3])).
fof(f3983, plain, (~ spl66_35 | ~ spl66_51), inference(avatar_split_clause, [], [f3979, f802, f734])).
fof(f3979, plain, (~ (e2 = op(e1, e3)) | ~ spl66_51), inference(backward_demodulation, [], [f253, f804])).
fof(f804, plain, ((e2 = op(e0, e3)) | ~ spl66_51), inference(avatar_component_clause, [], [f802])).
fof(f253, plain, ~ (op(e0, e3) = op(e1, e3)), inference(cnf_transformation, [], [f3])).
fof(f3976, plain, (~ spl66_49 | ~ spl66_53), inference(avatar_split_clause, [], [f3970, f811, f794])).
fof(f794, plain, (spl66_49 <=> (e0 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_49])])).
fof(f811, plain, (spl66_53 <=> (e0 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_53])])).
fof(f3970, plain, (~ (e0 = op(e0, e3)) | ~ spl66_53), inference(backward_demodulation, [], [f264, f813])).
fof(f813, plain, ((e0 = op(e0, e2)) | ~ spl66_53), inference(avatar_component_clause, [], [f811])).
fof(f264, plain, ~ (op(e0, e2) = op(e0, e3)), inference(cnf_transformation, [], [f3])).
fof(f3952, plain, (~ spl66_54 | ~ spl66_62), inference(avatar_split_clause, [], [f3951, f849, f815])).
fof(f815, plain, (spl66_54 <=> (e1 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_54])])).
fof(f3951, plain, (~ (e1 = op(e0, e2)) | ~ spl66_62), inference(forward_demodulation, [], [f260, f851])).
fof(f260, plain, ~ (op(e0, e0) = op(e0, e2)), inference(cnf_transformation, [], [f3])).
fof(f3941, plain, (~ spl66_34 | ~ spl66_42), inference(avatar_split_clause, [], [f3729, f764, f730])).
fof(f764, plain, (spl66_42 <=> (e1 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_42])])).
fof(f3729, plain, (~ (e1 = op(e1, e3)) | ~ spl66_42), inference(backward_demodulation, [], [f269, f766])).
fof(f766, plain, ((e1 = op(e1, e1)) | ~ spl66_42), inference(avatar_component_clause, [], [f764])).
fof(f269, plain, ~ (op(e1, e1) = op(e1, e3)), inference(cnf_transformation, [], [f3])).
fof(f3940, plain, (~ spl66_32 | ~ spl66_16), inference(avatar_split_clause, [], [f3939, f653, f721])).
fof(f721, plain, (spl66_32 <=> (e3 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_32])])).
fof(f3939, plain, (~ (e3 = op(e2, e0)) | ~ spl66_16), inference(forward_demodulation, [], [f240, f655])).
fof(f240, plain, ~ (op(e2, e0) = op(e3, e0)), inference(cnf_transformation, [], [f3])).
fof(f3934, plain, (~ spl66_26 | ~ spl66_42), inference(avatar_split_clause, [], [f3933, f764, f696])).
fof(f696, plain, (spl66_26 <=> (e1 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_26])])).
fof(f3933, plain, (~ (e1 = op(e2, e1)) | ~ spl66_42), inference(forward_demodulation, [], [f244, f766])).
fof(f244, plain, ~ (op(e1, e1) = op(e2, e1)), inference(cnf_transformation, [], [f3])).
fof(f3932, plain, (~ spl66_28 | ~ spl66_60), inference(avatar_split_clause, [], [f3931, f840, f704])).
fof(f840, plain, (spl66_60 <=> (e3 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_60])])).
fof(f3931, plain, (~ (e3 = op(e2, e1)) | ~ spl66_60), inference(forward_demodulation, [], [f242, f842])).
fof(f842, plain, ((e3 = op(e0, e1)) | ~ spl66_60), inference(avatar_component_clause, [], [f840])).
fof(f3929, plain, (~ spl66_12 | ~ spl66_16), inference(avatar_split_clause, [], [f3928, f653, f636])).
fof(f636, plain, (spl66_12 <=> (e3 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_12])])).
fof(f3928, plain, (~ (e3 = op(e3, e1)) | ~ spl66_16), inference(forward_demodulation, [], [f277, f655])).
fof(f3927, plain, (~ spl66_2 | ~ spl66_25 | ~ spl66_156 | spl66_165), inference(avatar_contradiction_clause, [], [f3926])).
fof(f3926, plain, ($false | (~ spl66_2 | ~ spl66_25 | ~ spl66_156 | spl66_165)), inference(subsumption_resolution, [], [f3923, f694])).
fof(f694, plain, ((e0 = op(e2, e1)) | ~ spl66_25), inference(avatar_component_clause, [], [f692])).
fof(f3923, plain, (~ (e0 = op(e2, e1)) | (~ spl66_2 | ~ spl66_156 | spl66_165)), inference(backward_demodulation, [], [f3820, f596])).
fof(f3820, plain, (~ (e0 = op(e2, op(e3, e3))) | (~ spl66_156 | spl66_165)), inference(backward_demodulation, [], [f1582, f1539])).
fof(f1539, plain, ((e2 = op(op(e3, e3), e3)) | ~ spl66_156), inference(avatar_component_clause, [], [f1538])).
fof(f1538, plain, (spl66_156 <=> (e2 = op(op(e3, e3), e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_156])])).
fof(f1582, plain, (~ (e0 = op(op(op(e3, e3), e3), op(e3, e3))) | spl66_165), inference(avatar_component_clause, [], [f1580])).
fof(f1580, plain, (spl66_165 <=> (e0 = op(op(op(e3, e3), e3), op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl66_165])])).
fof(f3916, plain, (~ spl66_1 | ~ spl66_5), inference(avatar_split_clause, [], [f3911, f607, f590])).
fof(f607, plain, (spl66_5 <=> (e0 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_5])])).
fof(f3911, plain, (~ (e0 = op(e3, e3)) | ~ spl66_5), inference(backward_demodulation, [], [f282, f609])).
fof(f609, plain, ((e0 = op(e3, e2)) | ~ spl66_5), inference(avatar_component_clause, [], [f607])).
fof(f282, plain, ~ (op(e3, e2) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f3910, plain, (~ spl66_1 | ~ spl66_16 | spl66_146), inference(avatar_split_clause, [], [f3908, f1489, f653, f590])).
fof(f1489, plain, (spl66_146 <=> (e0 = op(e3, op(e3, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl66_146])])).
fof(f3908, plain, (~ (e0 = op(e3, e3)) | (~ spl66_16 | spl66_146)), inference(backward_demodulation, [], [f1491, f655])).
fof(f1491, plain, (~ (e0 = op(e3, op(e3, e0))) | spl66_146), inference(avatar_component_clause, [], [f1489])).
fof(f3909, plain, (~ spl66_8 | ~ spl66_16), inference(avatar_split_clause, [], [f3906, f653, f619])).
fof(f619, plain, (spl66_8 <=> (e3 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_8])])).
fof(f3906, plain, (~ (e3 = op(e3, e2)) | ~ spl66_16), inference(backward_demodulation, [], [f278, f655])).
fof(f278, plain, ~ (op(e3, e0) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f3905, plain, (~ spl66_18 | ~ spl66_20), inference(avatar_contradiction_clause, [], [f3904])).
fof(f3904, plain, ($false | (~ spl66_18 | ~ spl66_20)), inference(subsumption_resolution, [], [f3903, f287])).
fof(f3903, plain, ((e1 = e3) | (~ spl66_18 | ~ spl66_20)), inference(backward_demodulation, [], [f672, f664])).
fof(f664, plain, ((e1 = op(e2, e3)) | ~ spl66_18), inference(avatar_component_clause, [], [f662])).
fof(f662, plain, (spl66_18 <=> (e1 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_18])])).
fof(f672, plain, ((e3 = op(e2, e3)) | ~ spl66_20), inference(avatar_component_clause, [], [f670])).
fof(f3894, plain, (~ spl66_19 | ~ spl66_23), inference(avatar_split_clause, [], [f3887, f683, f666])).
fof(f3887, plain, (~ (e2 = op(e2, e3)) | ~ spl66_23), inference(backward_demodulation, [], [f276, f685])).
fof(f685, plain, ((e2 = op(e2, e2)) | ~ spl66_23), inference(avatar_component_clause, [], [f683])).
fof(f276, plain, ~ (op(e2, e2) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f3885, plain, (~ spl66_19 | ~ spl66_35), inference(avatar_split_clause, [], [f3881, f734, f666])).
fof(f3881, plain, (~ (e2 = op(e2, e3)) | ~ spl66_35), inference(backward_demodulation, [], [f256, f736])).
fof(f736, plain, ((e2 = op(e1, e3)) | ~ spl66_35), inference(avatar_component_clause, [], [f734])).
fof(f256, plain, ~ (op(e1, e3) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f3880, plain, (spl66_35 | ~ spl66_40 | ~ spl66_135), inference(avatar_split_clause, [], [f3877, f1431, f755, f734])).
fof(f755, plain, (spl66_40 <=> (e3 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_40])])).
fof(f1431, plain, (spl66_135 <=> (e2 = op(e1, op(e1, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl66_135])])).
fof(f3877, plain, ((e2 = op(e1, e3)) | (~ spl66_40 | ~ spl66_135)), inference(backward_demodulation, [], [f1432, f757])).
fof(f757, plain, ((e3 = op(e1, e2)) | ~ spl66_40), inference(avatar_component_clause, [], [f755])).
fof(f1432, plain, ((e2 = op(e1, op(e1, e2))) | ~ spl66_135), inference(avatar_component_clause, [], [f1431])).
fof(f3879, plain, (~ spl66_36 | ~ spl66_40), inference(avatar_split_clause, [], [f3876, f755, f738])).
fof(f738, plain, (spl66_36 <=> (e3 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_36])])).
fof(f3876, plain, (~ (e3 = op(e1, e3)) | ~ spl66_40), inference(backward_demodulation, [], [f270, f757])).
fof(f270, plain, ~ (op(e1, e2) = op(e1, e3)), inference(cnf_transformation, [], [f3])).
fof(f3878, plain, (~ spl66_8 | ~ spl66_40), inference(avatar_split_clause, [], [f3875, f755, f619])).
fof(f3875, plain, (~ (e3 = op(e3, e2)) | ~ spl66_40), inference(backward_demodulation, [], [f251, f757])).
fof(f3873, plain, (~ spl66_37 | ~ spl66_45), inference(avatar_split_clause, [], [f3868, f777, f743])).
fof(f743, plain, (spl66_37 <=> (e0 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_37])])).
fof(f777, plain, (spl66_45 <=> (e0 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_45])])).
fof(f3868, plain, (~ (e0 = op(e1, e2)) | ~ spl66_45), inference(backward_demodulation, [], [f266, f779])).
fof(f779, plain, ((e0 = op(e1, e0)) | ~ spl66_45), inference(avatar_component_clause, [], [f777])).
fof(f266, plain, ~ (op(e1, e0) = op(e1, e2)), inference(cnf_transformation, [], [f3])).
fof(f3863, plain, (~ spl66_1 | ~ spl66_49), inference(avatar_split_clause, [], [f3858, f794, f590])).
fof(f3858, plain, (~ (e0 = op(e3, e3)) | ~ spl66_49), inference(backward_demodulation, [], [f255, f796])).
fof(f796, plain, ((e0 = op(e0, e3)) | ~ spl66_49), inference(avatar_component_clause, [], [f794])).
fof(f255, plain, ~ (op(e0, e3) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f3853, plain, (~ spl66_50 | ~ spl66_54), inference(avatar_split_clause, [], [f3848, f815, f798])).
fof(f798, plain, (spl66_50 <=> (e1 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_50])])).
fof(f3848, plain, (~ (e1 = op(e0, e3)) | ~ spl66_54), inference(backward_demodulation, [], [f264, f817])).
fof(f817, plain, ((e1 = op(e0, e2)) | ~ spl66_54), inference(avatar_component_clause, [], [f815])).
fof(f3839, plain, (spl66_53 | ~ spl66_63 | ~ spl66_142), inference(avatar_contradiction_clause, [], [f3838])).
fof(f3838, plain, ($false | (spl66_53 | ~ spl66_63 | ~ spl66_142)), inference(subsumption_resolution, [], [f3829, f812])).
fof(f812, plain, (~ (e0 = op(e0, e2)) | spl66_53), inference(avatar_component_clause, [], [f811])).
fof(f3829, plain, ((e0 = op(e0, e2)) | (~ spl66_63 | ~ spl66_142)), inference(backward_demodulation, [], [f1466, f855])).
fof(f855, plain, ((op(e0, e0) = e2) | ~ spl66_63), inference(avatar_component_clause, [], [f853])).
fof(f853, plain, (spl66_63 <=> (op(e0, e0) = e2)), introduced(avatar_definition, [new_symbols(naming, [spl66_63])])).
fof(f1466, plain, ((e0 = op(e0, op(e0, e0))) | ~ spl66_142), inference(avatar_component_clause, [], [f1465])).
fof(f1465, plain, (spl66_142 <=> (e0 = op(e0, op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl66_142])])).
fof(f3823, plain, (~ spl66_50 | ~ spl66_60 | spl66_141), inference(avatar_split_clause, [], [f3822, f1460, f840, f798])).
fof(f1460, plain, (spl66_141 <=> (e1 = op(e0, op(e0, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl66_141])])).
fof(f3822, plain, (~ (e1 = op(e0, e3)) | (~ spl66_60 | spl66_141)), inference(forward_demodulation, [], [f1462, f842])).
fof(f1462, plain, (~ (e1 = op(e0, op(e0, e1))) | spl66_141), inference(avatar_component_clause, [], [f1460])).
fof(f3796, plain, (~ spl66_48 | ~ spl66_32), inference(avatar_split_clause, [], [f3795, f721, f789])).
fof(f3795, plain, (~ (e3 = op(e1, e0)) | ~ spl66_32), inference(forward_demodulation, [], [f238, f723])).
fof(f723, plain, ((e3 = op(e2, e0)) | ~ spl66_32), inference(avatar_component_clause, [], [f721])).
fof(f238, plain, ~ (op(e1, e0) = op(e2, e0)), inference(cnf_transformation, [], [f3])).
fof(f3794, plain, (spl66_17 | ~ spl66_32 | ~ spl66_132), inference(avatar_contradiction_clause, [], [f3793])).
fof(f3793, plain, ($false | (spl66_17 | ~ spl66_32 | ~ spl66_132)), inference(subsumption_resolution, [], [f3792, f659])).
fof(f659, plain, (~ (e0 = op(e2, e3)) | spl66_17), inference(avatar_component_clause, [], [f658])).
fof(f658, plain, (spl66_17 <=> (e0 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_17])])).
fof(f3792, plain, ((e0 = op(e2, e3)) | (~ spl66_32 | ~ spl66_132)), inference(forward_demodulation, [], [f1418, f723])).
fof(f1418, plain, ((e0 = op(e2, op(e2, e0))) | ~ spl66_132), inference(avatar_component_clause, [], [f1417])).
fof(f1417, plain, (spl66_132 <=> (e0 = op(e2, op(e2, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl66_132])])).
fof(f3789, plain, (~ spl66_1 | ~ spl66_50 | spl66_158), inference(avatar_contradiction_clause, [], [f3788])).
fof(f3788, plain, ($false | (~ spl66_1 | ~ spl66_50 | spl66_158)), inference(subsumption_resolution, [], [f3784, f800])).
fof(f800, plain, ((e1 = op(e0, e3)) | ~ spl66_50), inference(avatar_component_clause, [], [f798])).
fof(f3784, plain, (~ (e1 = op(e0, e3)) | (~ spl66_1 | spl66_158)), inference(backward_demodulation, [], [f1549, f592])).
fof(f592, plain, ((e0 = op(e3, e3)) | ~ spl66_1), inference(avatar_component_clause, [], [f590])).
fof(f1549, plain, (~ (e1 = op(op(e3, e3), e3)) | spl66_158), inference(avatar_component_clause, [], [f1547])).
fof(f1547, plain, (spl66_158 <=> (e1 = op(op(e3, e3), e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_158])])).
fof(f3772, plain, (~ spl66_2 | ~ spl66_14), inference(avatar_split_clause, [], [f3768, f645, f594])).
fof(f3768, plain, (~ (e1 = op(e3, e3)) | ~ spl66_14), inference(backward_demodulation, [], [f279, f647])).
fof(f279, plain, ~ (op(e3, e0) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f3765, plain, (~ spl66_22 | ~ spl66_60 | spl66_153 | ~ spl66_164), inference(avatar_contradiction_clause, [], [f3764])).
fof(f3764, plain, ($false | (~ spl66_22 | ~ spl66_60 | spl66_153 | ~ spl66_164)), inference(subsumption_resolution, [], [f3762, f842])).
fof(f3762, plain, (~ (e3 = op(e0, e1)) | (~ spl66_22 | spl66_153 | ~ spl66_164)), inference(backward_demodulation, [], [f3702, f681])).
fof(f3702, plain, (~ (e3 = op(e0, op(e2, e2))) | (spl66_153 | ~ spl66_164)), inference(forward_demodulation, [], [f1527, f1576])).
fof(f1576, plain, ((e0 = op(op(e2, e2), e2)) | ~ spl66_164), inference(avatar_component_clause, [], [f1575])).
fof(f1575, plain, (spl66_164 <=> (e0 = op(op(e2, e2), e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_164])])).
fof(f1527, plain, (~ (e3 = op(op(op(e2, e2), e2), op(e2, e2))) | spl66_153), inference(avatar_component_clause, [], [f1525])).
fof(f1525, plain, (spl66_153 <=> (e3 = op(op(op(e2, e2), e2), op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl66_153])])).
fof(f3749, plain, (~ spl66_33 | ~ spl66_37), inference(avatar_split_clause, [], [f3743, f743, f726])).
fof(f3743, plain, (~ (e0 = op(e1, e3)) | ~ spl66_37), inference(backward_demodulation, [], [f270, f745])).
fof(f745, plain, ((e0 = op(e1, e2)) | ~ spl66_37), inference(avatar_component_clause, [], [f743])).
fof(f3736, plain, (~ spl66_38 | ~ spl66_42), inference(avatar_split_clause, [], [f3728, f764, f747])).
fof(f747, plain, (spl66_38 <=> (e1 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_38])])).
fof(f3728, plain, (~ (e1 = op(e1, e2)) | ~ spl66_42), inference(backward_demodulation, [], [f268, f766])).
fof(f268, plain, ~ (op(e1, e1) = op(e1, e2)), inference(cnf_transformation, [], [f3])).
fof(f3725, plain, (~ spl66_43 | ~ spl66_47), inference(avatar_split_clause, [], [f3719, f785, f768])).
fof(f768, plain, (spl66_43 <=> (e2 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_43])])).
fof(f3719, plain, (~ (e2 = op(e1, e1)) | ~ spl66_47), inference(backward_demodulation, [], [f265, f787])).
fof(f265, plain, ~ (op(e1, e0) = op(e1, e1)), inference(cnf_transformation, [], [f3])).
fof(f3724, plain, (~ spl66_15 | ~ spl66_47), inference(avatar_split_clause, [], [f3718, f785, f649])).
fof(f649, plain, (spl66_15 <=> (e2 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_15])])).
fof(f3718, plain, (~ (e2 = op(e3, e0)) | ~ spl66_47), inference(backward_demodulation, [], [f239, f787])).
fof(f3717, plain, (~ spl66_2 | ~ spl66_50), inference(avatar_split_clause, [], [f3713, f798, f594])).
fof(f3713, plain, (~ (e1 = op(e3, e3)) | ~ spl66_50), inference(backward_demodulation, [], [f255, f800])).
fof(f3710, plain, (~ spl66_52 | ~ spl66_60), inference(avatar_split_clause, [], [f3707, f840, f806])).
fof(f3707, plain, (~ (e3 = op(e0, e3)) | ~ spl66_60), inference(backward_demodulation, [], [f263, f842])).
fof(f3709, plain, (~ spl66_12 | ~ spl66_60), inference(avatar_split_clause, [], [f3706, f840, f636])).
fof(f3706, plain, (~ (e3 = op(e3, e1)) | ~ spl66_60), inference(backward_demodulation, [], [f243, f842])).
fof(f243, plain, ~ (op(e0, e1) = op(e3, e1)), inference(cnf_transformation, [], [f3])).
fof(f3693, plain, (~ spl66_14 | ~ spl66_30), inference(avatar_split_clause, [], [f3692, f713, f645])).
fof(f713, plain, (spl66_30 <=> (e1 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_30])])).
fof(f3692, plain, (~ (e1 = op(e3, e0)) | ~ spl66_30), inference(forward_demodulation, [], [f240, f715])).
fof(f715, plain, ((e1 = op(e2, e0)) | ~ spl66_30), inference(avatar_component_clause, [], [f713])).
fof(f3690, plain, (~ spl66_2 | ~ spl66_33 | spl66_166), inference(avatar_contradiction_clause, [], [f3689])).
fof(f3689, plain, ($false | (~ spl66_2 | ~ spl66_33 | spl66_166)), inference(subsumption_resolution, [], [f3682, f728])).
fof(f3682, plain, (~ (e0 = op(e1, e3)) | (~ spl66_2 | spl66_166)), inference(backward_demodulation, [], [f1587, f596])).
fof(f1587, plain, (~ (e0 = op(op(e3, e3), e3)) | spl66_166), inference(avatar_component_clause, [], [f1585])).
fof(f3668, plain, (~ spl66_4 | ~ spl66_12), inference(avatar_split_clause, [], [f3666, f636, f602])).
fof(f602, plain, (spl66_4 <=> (e3 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_4])])).
fof(f3666, plain, (~ (e3 = op(e3, e3)) | ~ spl66_12), inference(backward_demodulation, [], [f281, f638])).
fof(f638, plain, ((e3 = op(e3, e1)) | ~ spl66_12), inference(avatar_component_clause, [], [f636])).
fof(f3651, plain, (~ spl66_25 | ~ spl66_30 | spl66_131), inference(avatar_contradiction_clause, [], [f3650])).
fof(f3650, plain, ($false | (~ spl66_25 | ~ spl66_30 | spl66_131)), inference(subsumption_resolution, [], [f3647, f715])).
fof(f3647, plain, (~ (e1 = op(e2, e0)) | (~ spl66_25 | spl66_131)), inference(backward_demodulation, [], [f1414, f694])).
fof(f1414, plain, (~ (e1 = op(e2, op(e2, e1))) | spl66_131), inference(avatar_component_clause, [], [f1412])).
fof(f1412, plain, (spl66_131 <=> (e1 = op(e2, op(e2, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl66_131])])).
fof(f3649, plain, (~ spl66_17 | ~ spl66_25), inference(avatar_split_clause, [], [f3643, f692, f658])).
fof(f3643, plain, (~ (e0 = op(e2, e3)) | ~ spl66_25), inference(backward_demodulation, [], [f275, f694])).
fof(f3640, plain, (~ spl66_17 | ~ spl66_33), inference(avatar_split_clause, [], [f3634, f726, f658])).
fof(f3634, plain, (~ (e0 = op(e2, e3)) | ~ spl66_33), inference(backward_demodulation, [], [f256, f728])).
fof(f3633, plain, (~ spl66_38 | ~ spl66_43 | spl66_135), inference(avatar_contradiction_clause, [], [f3632])).
fof(f3632, plain, ($false | (~ spl66_38 | ~ spl66_43 | spl66_135)), inference(subsumption_resolution, [], [f3630, f770])).
fof(f770, plain, ((e2 = op(e1, e1)) | ~ spl66_43), inference(avatar_component_clause, [], [f768])).
fof(f3630, plain, (~ (e2 = op(e1, e1)) | (~ spl66_38 | spl66_135)), inference(backward_demodulation, [], [f1433, f749])).
fof(f749, plain, ((e1 = op(e1, e2)) | ~ spl66_38), inference(avatar_component_clause, [], [f747])).
fof(f1433, plain, (~ (e2 = op(e1, op(e1, e2))) | spl66_135), inference(avatar_component_clause, [], [f1431])).
fof(f3631, plain, (~ spl66_6 | ~ spl66_38), inference(avatar_split_clause, [], [f3626, f747, f611])).
fof(f611, plain, (spl66_6 <=> (e1 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_6])])).
fof(f3626, plain, (~ (e1 = op(e3, e2)) | ~ spl66_38), inference(backward_demodulation, [], [f251, f749])).
fof(f3619, plain, (~ spl66_35 | ~ spl66_43), inference(avatar_split_clause, [], [f3612, f768, f734])).
fof(f3612, plain, (~ (e2 = op(e1, e3)) | ~ spl66_43), inference(backward_demodulation, [], [f269, f770])).
fof(f3608, plain, (~ spl66_4 | ~ spl66_52), inference(avatar_split_clause, [], [f3606, f806, f602])).
fof(f3606, plain, (~ (e3 = op(e3, e3)) | ~ spl66_52), inference(backward_demodulation, [], [f255, f808])).
fof(f3603, plain, (~ spl66_23 | ~ spl66_55), inference(avatar_split_clause, [], [f3598, f819, f683])).
fof(f819, plain, (spl66_55 <=> (e2 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_55])])).
fof(f3598, plain, (~ (e2 = op(e2, e2)) | ~ spl66_55), inference(backward_demodulation, [], [f248, f821])).
fof(f821, plain, ((e2 = op(e0, e2)) | ~ spl66_55), inference(avatar_component_clause, [], [f819])).
fof(f248, plain, ~ (op(e0, e2) = op(e2, e2)), inference(cnf_transformation, [], [f3])).
fof(f3596, plain, (~ spl66_50 | ~ spl66_58), inference(avatar_split_clause, [], [f3591, f832, f798])).
fof(f832, plain, (spl66_58 <=> (e1 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_58])])).
fof(f3591, plain, (~ (e1 = op(e0, e3)) | ~ spl66_58), inference(backward_demodulation, [], [f263, f834])).
fof(f834, plain, ((e1 = op(e0, e1)) | ~ spl66_58), inference(avatar_component_clause, [], [f832])).
fof(f3595, plain, (~ spl66_42 | ~ spl66_58), inference(avatar_split_clause, [], [f3587, f832, f764])).
fof(f3587, plain, (~ (e1 = op(e1, e1)) | ~ spl66_58), inference(backward_demodulation, [], [f241, f834])).
fof(f241, plain, ~ (op(e0, e1) = op(e1, e1)), inference(cnf_transformation, [], [f3])).
fof(f3570, plain, (~ spl66_25 | ~ spl66_30 | spl66_132), inference(avatar_split_clause, [], [f3547, f1417, f713, f692])).
fof(f3547, plain, (~ (e0 = op(e2, e1)) | (~ spl66_30 | spl66_132)), inference(backward_demodulation, [], [f1419, f715])).
fof(f1419, plain, (~ (e0 = op(e2, op(e2, e0))) | spl66_132), inference(avatar_component_clause, [], [f1417])).
fof(f3553, plain, (~ spl66_18 | ~ spl66_28 | spl66_131), inference(avatar_split_clause, [], [f3551, f1412, f704, f662])).
fof(f3551, plain, (~ (e1 = op(e2, e3)) | (~ spl66_28 | spl66_131)), inference(backward_demodulation, [], [f1414, f706])).
fof(f3552, plain, (~ spl66_24 | ~ spl66_28), inference(avatar_split_clause, [], [f3549, f704, f687])).
fof(f687, plain, (spl66_24 <=> (e3 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_24])])).
fof(f3549, plain, (~ (e3 = op(e2, e2)) | ~ spl66_28), inference(backward_demodulation, [], [f274, f706])).
fof(f3525, plain, (~ spl66_42 | spl66_136), inference(avatar_contradiction_clause, [], [f3524])).
fof(f3524, plain, ($false | (~ spl66_42 | spl66_136)), inference(subsumption_resolution, [], [f3519, f766])).
fof(f3519, plain, (~ (e1 = op(e1, e1)) | (~ spl66_42 | spl66_136)), inference(backward_demodulation, [], [f1438, f766])).
fof(f1438, plain, (~ (e1 = op(e1, op(e1, e1))) | spl66_136), inference(avatar_component_clause, [], [f1436])).
fof(f1436, plain, (spl66_136 <=> (e1 = op(e1, op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl66_136])])).
fof(f3513, plain, (~ spl66_33 | ~ spl66_48 | spl66_137), inference(avatar_split_clause, [], [f3511, f1441, f789, f726])).
fof(f1441, plain, (spl66_137 <=> (e0 = op(e1, op(e1, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl66_137])])).
fof(f3511, plain, (~ (e0 = op(e1, e3)) | (~ spl66_48 | spl66_137)), inference(backward_demodulation, [], [f1443, f791])).
fof(f1443, plain, (~ (e0 = op(e1, op(e1, e0))) | spl66_137), inference(avatar_component_clause, [], [f1441])).
fof(f3512, plain, (~ spl66_44 | ~ spl66_48), inference(avatar_split_clause, [], [f3508, f789, f772])).
fof(f772, plain, (spl66_44 <=> (e3 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_44])])).
fof(f3508, plain, (~ (e3 = op(e1, e1)) | ~ spl66_48), inference(backward_demodulation, [], [f265, f791])).
fof(f3506, plain, (~ spl66_18 | ~ spl66_50), inference(avatar_split_clause, [], [f3503, f798, f662])).
fof(f3503, plain, (~ (e1 = op(e2, e3)) | ~ spl66_50), inference(backward_demodulation, [], [f254, f800])).
fof(f3501, plain, (~ spl66_24 | ~ spl66_56), inference(avatar_split_clause, [], [f3499, f823, f687])).
fof(f823, plain, (spl66_56 <=> (e3 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_56])])).
fof(f3499, plain, (~ (e3 = op(e2, e2)) | ~ spl66_56), inference(backward_demodulation, [], [f248, f825])).
fof(f825, plain, ((e3 = op(e0, e2)) | ~ spl66_56), inference(avatar_component_clause, [], [f823])).
fof(f3481, plain, (~ spl66_53 | ~ spl66_61), inference(avatar_split_clause, [], [f3470, f845, f811])).
fof(f845, plain, (spl66_61 <=> (e0 = op(e0, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_61])])).
fof(f3470, plain, (~ (e0 = op(e0, e2)) | ~ spl66_61), inference(backward_demodulation, [], [f260, f847])).
fof(f847, plain, ((e0 = op(e0, e0)) | ~ spl66_61), inference(avatar_component_clause, [], [f845])).
fof(f3480, plain, (~ spl66_29 | ~ spl66_61), inference(avatar_split_clause, [], [f3468, f845, f709])).
fof(f3468, plain, (~ (e0 = op(e2, e0)) | ~ spl66_61), inference(backward_demodulation, [], [f236, f847])).
fof(f236, plain, ~ (op(e0, e0) = op(e2, e0)), inference(cnf_transformation, [], [f3])).
fof(f3444, plain, (~ spl66_7 | ~ spl66_15), inference(avatar_split_clause, [], [f3441, f649, f615])).
fof(f3441, plain, (~ (e2 = op(e3, e2)) | ~ spl66_15), inference(backward_demodulation, [], [f278, f651])).
fof(f651, plain, ((e2 = op(e3, e0)) | ~ spl66_15), inference(avatar_component_clause, [], [f649])).
fof(f3437, plain, (~ spl66_24 | ~ spl66_33 | ~ spl66_154 | spl66_163), inference(avatar_contradiction_clause, [], [f3436])).
fof(f3436, plain, ($false | (~ spl66_24 | ~ spl66_33 | ~ spl66_154 | spl66_163)), inference(subsumption_resolution, [], [f3434, f728])).
fof(f3434, plain, (~ (e0 = op(e1, e3)) | (~ spl66_24 | ~ spl66_154 | spl66_163)), inference(backward_demodulation, [], [f3356, f689])).
fof(f689, plain, ((e3 = op(e2, e2)) | ~ spl66_24), inference(avatar_component_clause, [], [f687])).
fof(f3356, plain, (~ (e0 = op(e1, op(e2, e2))) | (~ spl66_154 | spl66_163)), inference(backward_demodulation, [], [f1572, f1530])).
fof(f1530, plain, ((e1 = op(op(e2, e2), e2)) | ~ spl66_154), inference(avatar_component_clause, [], [f1529])).
fof(f1529, plain, (spl66_154 <=> (e1 = op(op(e2, e2), e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_154])])).
fof(f3428, plain, (~ spl66_17 | ~ spl66_29), inference(avatar_split_clause, [], [f3422, f709, f658])).
fof(f3422, plain, (~ (e0 = op(e2, e3)) | ~ spl66_29), inference(backward_demodulation, [], [f273, f711])).
fof(f3427, plain, (~ spl66_13 | ~ spl66_29), inference(avatar_split_clause, [], [f3420, f709, f641])).
fof(f3420, plain, (~ (e0 = op(e3, e0)) | ~ spl66_29), inference(backward_demodulation, [], [f240, f711])).
fof(f3412, plain, (~ spl66_39 | spl66_135), inference(avatar_contradiction_clause, [], [f3411])).
fof(f3411, plain, ($false | (~ spl66_39 | spl66_135)), inference(subsumption_resolution, [], [f3408, f753])).
fof(f3408, plain, (~ (e2 = op(e1, e2)) | (~ spl66_39 | spl66_135)), inference(backward_demodulation, [], [f1433, f753])).
fof(f3410, plain, (~ spl66_35 | ~ spl66_39), inference(avatar_split_clause, [], [f3406, f751, f734])).
fof(f3406, plain, (~ (e2 = op(e1, e3)) | ~ spl66_39), inference(backward_demodulation, [], [f270, f753])).
fof(f3401, plain, (~ spl66_9 | ~ spl66_44 | spl66_170), inference(avatar_split_clause, [], [f3395, f1605, f772, f624])).
fof(f1605, plain, (spl66_170 <=> (e0 = op(op(e1, e1), e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_170])])).
fof(f3395, plain, (~ (e0 = op(e3, e1)) | (~ spl66_44 | spl66_170)), inference(backward_demodulation, [], [f1607, f774])).
fof(f774, plain, ((e3 = op(e1, e1)) | ~ spl66_44), inference(avatar_component_clause, [], [f772])).
fof(f1607, plain, (~ (e0 = op(op(e1, e1), e1)) | spl66_170), inference(avatar_component_clause, [], [f1605])).
fof(f3400, plain, (spl66_12 | ~ spl66_44 | ~ spl66_148), inference(avatar_contradiction_clause, [], [f3399])).
fof(f3399, plain, ($false | (spl66_12 | ~ spl66_44 | ~ spl66_148)), inference(subsumption_resolution, [], [f3393, f637])).
fof(f637, plain, (~ (e3 = op(e3, e1)) | spl66_12), inference(avatar_component_clause, [], [f636])).
fof(f3393, plain, ((e3 = op(e3, e1)) | (~ spl66_44 | ~ spl66_148)), inference(backward_demodulation, [], [f1503, f774])).
fof(f1503, plain, ((e3 = op(op(e1, e1), e1)) | ~ spl66_148), inference(avatar_component_clause, [], [f1502])).
fof(f1502, plain, (spl66_148 <=> (e3 = op(op(e1, e1), e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_148])])).
fof(f3383, plain, (~ spl66_10 | ~ spl66_58), inference(avatar_split_clause, [], [f3378, f832, f628])).
fof(f628, plain, (spl66_10 <=> (e1 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_10])])).
fof(f3378, plain, (~ (e1 = op(e3, e1)) | ~ spl66_58), inference(backward_demodulation, [], [f243, f834])).
fof(f3374, plain, (~ spl66_15 | ~ spl66_64 | spl66_162), inference(avatar_split_clause, [], [f3367, f1565, f857, f649])).
fof(f857, plain, (spl66_64 <=> (op(e0, e0) = e3)), introduced(avatar_definition, [new_symbols(naming, [spl66_64])])).
fof(f1565, plain, (spl66_162 <=> (e2 = op(op(e0, e0), e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_162])])).
fof(f3367, plain, (~ (e2 = op(e3, e0)) | (~ spl66_64 | spl66_162)), inference(backward_demodulation, [], [f1567, f859])).
fof(f859, plain, ((op(e0, e0) = e3) | ~ spl66_64), inference(avatar_component_clause, [], [f857])).
fof(f1567, plain, (~ (e2 = op(op(e0, e0), e0)) | spl66_162), inference(avatar_component_clause, [], [f1565])).
fof(f3371, plain, (~ spl66_60 | ~ spl66_64), inference(avatar_split_clause, [], [f3363, f857, f840])).
fof(f3363, plain, (~ (e3 = op(e0, e1)) | ~ spl66_64), inference(backward_demodulation, [], [f259, f859])).
fof(f259, plain, ~ (op(e0, e0) = op(e0, e1)), inference(cnf_transformation, [], [f3])).
fof(f3370, plain, (~ spl66_32 | ~ spl66_64), inference(avatar_split_clause, [], [f3361, f857, f721])).
fof(f3361, plain, (~ (e3 = op(e2, e0)) | ~ spl66_64), inference(backward_demodulation, [], [f236, f859])).
fof(f3360, plain, (~ spl66_22 | ~ spl66_27 | spl66_131), inference(avatar_split_clause, [], [f3359, f1412, f700, f679])).
fof(f3359, plain, (~ (e1 = op(e2, e2)) | (~ spl66_27 | spl66_131)), inference(forward_demodulation, [], [f1414, f702])).
fof(f3358, plain, (~ spl66_41 | ~ spl66_46 | spl66_137), inference(avatar_split_clause, [], [f3357, f1441, f781, f760])).
fof(f760, plain, (spl66_41 <=> (e0 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_41])])).
fof(f781, plain, (spl66_46 <=> (e1 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_46])])).
fof(f3357, plain, (~ (e0 = op(e1, e1)) | (~ spl66_46 | spl66_137)), inference(forward_demodulation, [], [f1443, f783])).
fof(f783, plain, ((e1 = op(e1, e0)) | ~ spl66_46), inference(avatar_component_clause, [], [f781])).
fof(f3351, plain, (~ spl66_63 | ~ spl66_53 | spl66_140), inference(avatar_split_clause, [], [f3298, f1455, f811, f853])).
fof(f1455, plain, (spl66_140 <=> (e2 = op(e0, op(e0, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl66_140])])).
fof(f3298, plain, (~ (op(e0, e0) = e2) | (~ spl66_53 | spl66_140)), inference(backward_demodulation, [], [f1457, f813])).
fof(f1457, plain, (~ (e2 = op(e0, op(e0, e2))) | spl66_140), inference(avatar_component_clause, [], [f1455])).
fof(f3345, plain, (~ spl66_42 | ~ spl66_46), inference(avatar_split_clause, [], [f3344, f781, f764])).
fof(f3344, plain, (~ (e1 = op(e1, e1)) | ~ spl66_46), inference(forward_demodulation, [], [f265, f783])).
fof(f3335, plain, (~ spl66_5 | ~ spl66_13), inference(avatar_split_clause, [], [f3330, f641, f607])).
fof(f3330, plain, (~ (e0 = op(e3, e2)) | ~ spl66_13), inference(backward_demodulation, [], [f278, f643])).
fof(f3328, plain, (~ spl66_22 | ~ spl66_40 | spl66_152), inference(avatar_contradiction_clause, [], [f3327])).
fof(f3327, plain, ($false | (~ spl66_22 | ~ spl66_40 | spl66_152)), inference(subsumption_resolution, [], [f3322, f757])).
fof(f3322, plain, (~ (e3 = op(e1, e2)) | (~ spl66_22 | spl66_152)), inference(backward_demodulation, [], [f1522, f681])).
fof(f1522, plain, (~ (e3 = op(op(e2, e2), e2)) | spl66_152), inference(avatar_component_clause, [], [f1520])).
fof(f3309, plain, (~ spl66_38 | ~ spl66_46), inference(avatar_split_clause, [], [f3304, f781, f747])).
fof(f3304, plain, (~ (e1 = op(e1, e2)) | ~ spl66_46), inference(backward_demodulation, [], [f266, f783])).
fof(f3308, plain, (~ spl66_30 | ~ spl66_46), inference(avatar_split_clause, [], [f3302, f781, f713])).
fof(f3302, plain, (~ (e1 = op(e2, e0)) | ~ spl66_46), inference(backward_demodulation, [], [f238, f783])).
fof(f3287, plain, (~ spl66_15 | ~ spl66_63), inference(avatar_split_clause, [], [f3279, f853, f649])).
fof(f3279, plain, (~ (e2 = op(e3, e0)) | ~ spl66_63), inference(backward_demodulation, [], [f237, f855])).
fof(f3272, plain, (spl66_46 | ~ spl66_41 | ~ spl66_136), inference(avatar_split_clause, [], [f3271, f1436, f760, f781])).
fof(f3271, plain, ((e1 = op(e1, e0)) | (~ spl66_41 | ~ spl66_136)), inference(forward_demodulation, [], [f1437, f762])).
fof(f762, plain, ((e0 = op(e1, e1)) | ~ spl66_41), inference(avatar_component_clause, [], [f760])).
fof(f1437, plain, ((e1 = op(e1, op(e1, e1))) | ~ spl66_136), inference(avatar_component_clause, [], [f1436])).
fof(f3246, plain, (~ spl66_15 | ~ spl66_41 | ~ spl66_60 | spl66_147), inference(avatar_split_clause, [], [f3163, f1498, f840, f760, f649])).
fof(f1498, plain, (spl66_147 <=> (e2 = op(op(op(e1, e1), e1), op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl66_147])])).
fof(f3163, plain, (~ (e2 = op(e3, e0)) | (~ spl66_41 | ~ spl66_60 | spl66_147)), inference(forward_demodulation, [], [f3162, f842])).
fof(f3162, plain, (~ (e2 = op(op(e0, e1), e0)) | (~ spl66_41 | spl66_147)), inference(forward_demodulation, [], [f1500, f762])).
fof(f1500, plain, (~ (e2 = op(op(op(e1, e1), e1), op(e1, e1))) | spl66_147), inference(avatar_component_clause, [], [f1498])).
fof(f3223, plain, (~ spl66_1 | ~ spl66_17), inference(avatar_split_clause, [], [f3218, f658, f590])).
fof(f3218, plain, (~ (e0 = op(e3, e3)) | ~ spl66_17), inference(backward_demodulation, [], [f258, f660])).
fof(f660, plain, ((e0 = op(e2, e3)) | ~ spl66_17), inference(avatar_component_clause, [], [f658])).
fof(f258, plain, ~ (op(e2, e3) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f3217, plain, (~ spl66_24 | ~ spl66_50 | spl66_151 | ~ spl66_164), inference(avatar_contradiction_clause, [], [f3216])).
fof(f3216, plain, ($false | (~ spl66_24 | ~ spl66_50 | spl66_151 | ~ spl66_164)), inference(subsumption_resolution, [], [f3212, f800])).
fof(f3212, plain, (~ (e1 = op(e0, e3)) | (~ spl66_24 | spl66_151 | ~ spl66_164)), inference(backward_demodulation, [], [f3171, f689])).
fof(f3171, plain, (~ (e1 = op(e0, op(e2, e2))) | (spl66_151 | ~ spl66_164)), inference(forward_demodulation, [], [f1518, f1576])).
fof(f1518, plain, (~ (e1 = op(op(op(e2, e2), e2), op(e2, e2))) | spl66_151), inference(avatar_component_clause, [], [f1516])).
fof(f1516, plain, (spl66_151 <=> (e1 = op(op(op(e2, e2), e2), op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl66_151])])).
fof(f3213, plain, (~ spl66_8 | ~ spl66_24), inference(avatar_split_clause, [], [f3207, f687, f619])).
fof(f3207, plain, (~ (e3 = op(e3, e2)) | ~ spl66_24), inference(backward_demodulation, [], [f252, f689])).
fof(f252, plain, ~ (op(e2, e2) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f3200, plain, (~ spl66_34 | ~ spl66_38), inference(avatar_split_clause, [], [f3196, f747, f730])).
fof(f3196, plain, (~ (e1 = op(e1, e3)) | ~ spl66_38), inference(backward_demodulation, [], [f270, f749])).
fof(f3193, plain, (~ spl66_34 | ~ spl66_50), inference(avatar_split_clause, [], [f3190, f798, f730])).
fof(f3190, plain, (~ (e1 = op(e1, e3)) | ~ spl66_50), inference(backward_demodulation, [], [f253, f800])).
fof(f3187, plain, (~ spl66_51 | ~ spl66_55), inference(avatar_split_clause, [], [f3183, f819, f802])).
fof(f3183, plain, (~ (e2 = op(e0, e3)) | ~ spl66_55), inference(backward_demodulation, [], [f264, f821])).
fof(f3186, plain, (~ spl66_39 | ~ spl66_55), inference(avatar_split_clause, [], [f3180, f819, f751])).
fof(f3180, plain, (~ (e2 = op(e1, e2)) | ~ spl66_55), inference(backward_demodulation, [], [f247, f821])).
fof(f247, plain, ~ (op(e0, e2) = op(e1, e2)), inference(cnf_transformation, [], [f3])).
fof(f3137, plain, (~ spl66_2 | ~ spl66_10), inference(avatar_split_clause, [], [f3133, f628, f594])).
fof(f3133, plain, (~ (e1 = op(e3, e3)) | ~ spl66_10), inference(backward_demodulation, [], [f281, f630])).
fof(f630, plain, ((e1 = op(e3, e1)) | ~ spl66_10), inference(avatar_component_clause, [], [f628])).
fof(f3088, plain, (~ spl66_23 | ~ spl66_39), inference(avatar_split_clause, [], [f3083, f751, f683])).
fof(f3083, plain, (~ (e2 = op(e2, e2)) | ~ spl66_39), inference(backward_demodulation, [], [f250, f753])).
fof(f250, plain, ~ (op(e1, e2) = op(e2, e2)), inference(cnf_transformation, [], [f3])).
fof(f3059, plain, (~ spl66_49 | ~ spl66_61), inference(avatar_split_clause, [], [f3049, f845, f794])).
fof(f3049, plain, (~ (e0 = op(e0, e3)) | ~ spl66_61), inference(backward_demodulation, [], [f261, f847])).
fof(f261, plain, ~ (op(e0, e0) = op(e0, e3)), inference(cnf_transformation, [], [f3])).
fof(f3045, plain, (~ spl66_5 | ~ spl66_15 | spl66_146), inference(avatar_split_clause, [], [f3044, f1489, f649, f607])).
fof(f3044, plain, (~ (e0 = op(e3, e2)) | (~ spl66_15 | spl66_146)), inference(forward_demodulation, [], [f1491, f651])).
fof(f3043, plain, (spl66_60 | ~ spl66_41 | ~ spl66_148), inference(avatar_split_clause, [], [f3042, f1502, f760, f840])).
fof(f3042, plain, ((e3 = op(e0, e1)) | (~ spl66_41 | ~ spl66_148)), inference(forward_demodulation, [], [f1503, f762])).
fof(f3004, plain, (~ spl66_8 | ~ spl66_12), inference(avatar_split_clause, [], [f3001, f636, f619])).
fof(f3001, plain, (~ (e3 = op(e3, e2)) | ~ spl66_12), inference(backward_demodulation, [], [f280, f638])).
fof(f2999, plain, (~ spl66_3 | ~ spl66_15), inference(avatar_split_clause, [], [f2994, f649, f598])).
fof(f598, plain, (spl66_3 <=> (e2 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_3])])).
fof(f2994, plain, (~ (e2 = op(e3, e3)) | ~ spl66_15), inference(backward_demodulation, [], [f279, f651])).
fof(f2977, plain, (~ spl66_18 | ~ spl66_26), inference(avatar_split_clause, [], [f2973, f696, f662])).
fof(f2973, plain, (~ (e1 = op(e2, e3)) | ~ spl66_26), inference(backward_demodulation, [], [f275, f698])).
fof(f698, plain, ((e1 = op(e2, e1)) | ~ spl66_26), inference(avatar_component_clause, [], [f696])).
fof(f2949, plain, (~ spl66_32 | ~ spl66_41 | spl66_149 | ~ spl66_150), inference(avatar_split_clause, [], [f2946, f1511, f1507, f760, f721])).
fof(f1507, plain, (spl66_149 <=> (e3 = op(op(op(e1, e1), e1), op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl66_149])])).
fof(f1511, plain, (spl66_150 <=> (e2 = op(op(e1, e1), e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_150])])).
fof(f2946, plain, (~ (e3 = op(e2, e0)) | (~ spl66_41 | spl66_149 | ~ spl66_150)), inference(backward_demodulation, [], [f2899, f762])).
fof(f2899, plain, (~ (e3 = op(e2, op(e1, e1))) | (spl66_149 | ~ spl66_150)), inference(forward_demodulation, [], [f1509, f1512])).
fof(f1512, plain, ((e2 = op(op(e1, e1), e1)) | ~ spl66_150), inference(avatar_component_clause, [], [f1511])).
fof(f1509, plain, (~ (e3 = op(op(op(e1, e1), e1), op(e1, e1))) | spl66_149), inference(avatar_component_clause, [], [f1507])).
fof(f2948, plain, (~ spl66_33 | ~ spl66_41), inference(avatar_split_clause, [], [f2938, f760, f726])).
fof(f2938, plain, (~ (e0 = op(e1, e3)) | ~ spl66_41), inference(backward_demodulation, [], [f269, f762])).
fof(f2947, plain, (~ spl66_9 | ~ spl66_41), inference(avatar_split_clause, [], [f2936, f760, f624])).
fof(f2936, plain, (~ (e0 = op(e3, e1)) | ~ spl66_41), inference(backward_demodulation, [], [f245, f762])).
fof(f245, plain, ~ (op(e1, e1) = op(e3, e1)), inference(cnf_transformation, [], [f3])).
fof(f2931, plain, (~ spl66_6 | ~ spl66_54), inference(avatar_split_clause, [], [f2927, f815, f611])).
fof(f2927, plain, (~ (e1 = op(e3, e2)) | ~ spl66_54), inference(backward_demodulation, [], [f249, f817])).
fof(f249, plain, ~ (op(e0, e2) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f2923, plain, (~ spl66_55 | ~ spl66_59), inference(avatar_split_clause, [], [f2919, f836, f819])).
fof(f2919, plain, (~ (e2 = op(e0, e2)) | ~ spl66_59), inference(backward_demodulation, [], [f262, f838])).
fof(f262, plain, ~ (op(e0, e1) = op(e0, e2)), inference(cnf_transformation, [], [f3])).
fof(f2915, plain, (~ spl66_49 | ~ spl66_64 | spl66_142), inference(avatar_contradiction_clause, [], [f2914])).
fof(f2914, plain, ($false | (~ spl66_49 | ~ spl66_64 | spl66_142)), inference(subsumption_resolution, [], [f2908, f796])).
fof(f2908, plain, (~ (e0 = op(e0, e3)) | (~ spl66_64 | spl66_142)), inference(backward_demodulation, [], [f1467, f859])).
fof(f1467, plain, (~ (e0 = op(e0, op(e0, e0))) | spl66_142), inference(avatar_component_clause, [], [f1465])).
fof(f2903, plain, (spl66_41 | ~ spl66_46 | ~ spl66_137), inference(avatar_split_clause, [], [f2902, f1441, f781, f760])).
fof(f2902, plain, ((e0 = op(e1, e1)) | (~ spl66_46 | ~ spl66_137)), inference(forward_demodulation, [], [f1442, f783])).
fof(f1442, plain, ((e0 = op(e1, op(e1, e0))) | ~ spl66_137), inference(avatar_component_clause, [], [f1441])).
fof(f2890, plain, (~ spl66_33 | ~ spl66_49), inference(avatar_split_clause, [], [f2889, f794, f726])).
fof(f2889, plain, (~ (e0 = op(e1, e3)) | ~ spl66_49), inference(forward_demodulation, [], [f253, f796])).
fof(f2883, plain, (~ spl66_6 | ~ spl66_8), inference(avatar_contradiction_clause, [], [f2882])).
fof(f2882, plain, ($false | (~ spl66_6 | ~ spl66_8)), inference(subsumption_resolution, [], [f2881, f287])).
fof(f2881, plain, ((e1 = e3) | (~ spl66_6 | ~ spl66_8)), inference(backward_demodulation, [], [f621, f613])).
fof(f613, plain, ((e1 = op(e3, e2)) | ~ spl66_6), inference(avatar_component_clause, [], [f611])).
fof(f621, plain, ((e3 = op(e3, e2)) | ~ spl66_8), inference(avatar_component_clause, [], [f619])).
fof(f2864, plain, (~ spl66_17 | ~ spl66_21), inference(avatar_split_clause, [], [f2855, f675, f658])).
fof(f675, plain, (spl66_21 <=> (e0 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_21])])).
fof(f2855, plain, (~ (e0 = op(e2, e3)) | ~ spl66_21), inference(backward_demodulation, [], [f276, f677])).
fof(f677, plain, ((e0 = op(e2, e2)) | ~ spl66_21), inference(avatar_component_clause, [], [f675])).
fof(f2842, plain, (~ spl66_34 | ~ spl66_46), inference(avatar_split_clause, [], [f2839, f781, f730])).
fof(f2839, plain, (~ (e1 = op(e1, e3)) | ~ spl66_46), inference(backward_demodulation, [], [f267, f783])).
fof(f2841, plain, (~ spl66_14 | ~ spl66_46), inference(avatar_split_clause, [], [f2838, f781, f645])).
fof(f2838, plain, (~ (e1 = op(e3, e0)) | ~ spl66_46), inference(backward_demodulation, [], [f239, f783])).
fof(f2835, plain, (~ spl66_49 | ~ spl66_52), inference(avatar_contradiction_clause, [], [f2834])).
fof(f2834, plain, ($false | (~ spl66_49 | ~ spl66_52)), inference(subsumption_resolution, [], [f2832, f285])).
fof(f285, plain, ~ (e0 = e3), inference(cnf_transformation, [], [f4])).
fof(f2832, plain, ((e0 = e3) | (~ spl66_49 | ~ spl66_52)), inference(backward_demodulation, [], [f808, f796])).
fof(f2829, plain, (~ spl66_6 | ~ spl66_32 | ~ spl66_63 | spl66_167), inference(avatar_split_clause, [], [f2828, f1590, f853, f721, f611])).
fof(f1590, plain, (spl66_167 <=> (e1 = op(op(op(e0, e0), e0), op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl66_167])])).
fof(f2828, plain, (~ (e1 = op(e3, e2)) | (~ spl66_32 | ~ spl66_63 | spl66_167)), inference(forward_demodulation, [], [f2820, f723])).
fof(f2820, plain, (~ (e1 = op(op(e2, e0), e2)) | (~ spl66_63 | spl66_167)), inference(backward_demodulation, [], [f1592, f855])).
fof(f1592, plain, (~ (e1 = op(op(op(e0, e0), e0), op(e0, e0))) | spl66_167), inference(avatar_component_clause, [], [f1590])).
fof(f2824, plain, (~ spl66_32 | ~ spl66_63 | spl66_160), inference(avatar_contradiction_clause, [], [f2823])).
fof(f2823, plain, ($false | (~ spl66_32 | ~ spl66_63 | spl66_160)), inference(subsumption_resolution, [], [f2818, f723])).
fof(f2818, plain, (~ (e3 = op(e2, e0)) | (~ spl66_63 | spl66_160)), inference(backward_demodulation, [], [f1558, f855])).
fof(f1558, plain, (~ (e3 = op(op(e0, e0), e0)) | spl66_160), inference(avatar_component_clause, [], [f1556])).
fof(f1556, plain, (spl66_160 <=> (e3 = op(op(e0, e0), e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_160])])).
fof(f2814, plain, (~ spl66_17 | ~ spl66_32 | spl66_132), inference(avatar_split_clause, [], [f2813, f1417, f721, f658])).
fof(f2813, plain, (~ (e0 = op(e2, e3)) | (~ spl66_32 | spl66_132)), inference(forward_demodulation, [], [f1419, f723])).
fof(f2812, plain, (~ spl66_34 | ~ spl66_44 | spl66_136), inference(avatar_split_clause, [], [f2811, f1436, f772, f730])).
fof(f2811, plain, (~ (e1 = op(e1, e3)) | (~ spl66_44 | spl66_136)), inference(forward_demodulation, [], [f1438, f774])).
fof(f2782, plain, (~ spl66_19 | ~ spl66_3), inference(avatar_split_clause, [], [f2781, f598, f666])).
fof(f2781, plain, (~ (e2 = op(e2, e3)) | ~ spl66_3), inference(forward_demodulation, [], [f258, f600])).
fof(f600, plain, ((e2 = op(e3, e3)) | ~ spl66_3), inference(avatar_component_clause, [], [f598])).
fof(f2775, plain, (~ spl66_3 | ~ spl66_17 | spl66_166), inference(avatar_contradiction_clause, [], [f2774])).
fof(f2774, plain, ($false | (~ spl66_3 | ~ spl66_17 | spl66_166)), inference(subsumption_resolution, [], [f2769, f660])).
fof(f2769, plain, (~ (e0 = op(e2, e3)) | (~ spl66_3 | spl66_166)), inference(backward_demodulation, [], [f1587, f600])).
fof(f2762, plain, (~ spl66_4 | ~ spl66_8), inference(avatar_split_clause, [], [f2760, f619, f602])).
fof(f2760, plain, (~ (e3 = op(e3, e3)) | ~ spl66_8), inference(backward_demodulation, [], [f282, f621])).
fof(f2754, plain, (~ spl66_22 | ~ spl66_37 | spl66_164), inference(avatar_contradiction_clause, [], [f2753])).
fof(f2753, plain, ($false | (~ spl66_22 | ~ spl66_37 | spl66_164)), inference(subsumption_resolution, [], [f2747, f745])).
fof(f2747, plain, (~ (e0 = op(e1, e2)) | (~ spl66_22 | spl66_164)), inference(backward_demodulation, [], [f1577, f681])).
fof(f1577, plain, (~ (e0 = op(op(e2, e2), e2)) | spl66_164), inference(avatar_component_clause, [], [f1575])).
fof(f2748, plain, (~ spl66_18 | ~ spl66_22), inference(avatar_split_clause, [], [f2740, f679, f662])).
fof(f2740, plain, (~ (e1 = op(e2, e3)) | ~ spl66_22), inference(backward_demodulation, [], [f276, f681])).
fof(f2738, plain, (~ spl66_18 | ~ spl66_34), inference(avatar_split_clause, [], [f2733, f730, f662])).
fof(f2733, plain, (~ (e1 = op(e2, e3)) | ~ spl66_34), inference(backward_demodulation, [], [f256, f732])).
fof(f2729, plain, (~ spl66_21 | ~ spl66_37), inference(avatar_split_clause, [], [f2722, f743, f675])).
fof(f2722, plain, (~ (e0 = op(e2, e2)) | ~ spl66_37), inference(backward_demodulation, [], [f250, f745])).
fof(f2709, plain, (~ spl66_51 | ~ spl66_56 | spl66_140), inference(avatar_split_clause, [], [f2708, f1455, f823, f802])).
fof(f2708, plain, (~ (e2 = op(e0, e3)) | (~ spl66_56 | spl66_140)), inference(forward_demodulation, [], [f1457, f825])).
fof(f2704, plain, (~ spl66_61 | spl66_142), inference(avatar_contradiction_clause, [], [f2703])).
fof(f2703, plain, ($false | (~ spl66_61 | spl66_142)), inference(subsumption_resolution, [], [f2702, f847])).
fof(f2702, plain, (~ (e0 = op(e0, e0)) | (~ spl66_61 | spl66_142)), inference(forward_demodulation, [], [f1467, f847])).
fof(f2698, plain, (~ spl66_52 | ~ spl66_56), inference(avatar_split_clause, [], [f2657, f823, f806])).
fof(f2657, plain, (~ (e3 = op(e0, e3)) | ~ spl66_56), inference(backward_demodulation, [], [f264, f825])).
fof(f2697, plain, (~ spl66_39 | ~ spl66_47), inference(avatar_split_clause, [], [f2696, f785, f751])).
fof(f2696, plain, (~ (e2 = op(e1, e2)) | ~ spl66_47), inference(forward_demodulation, [], [f266, f787])).
fof(f2693, plain, (~ spl66_37 | ~ spl66_47 | spl66_137), inference(avatar_split_clause, [], [f2661, f1441, f785, f743])).
fof(f2661, plain, (~ (e0 = op(e1, e2)) | (~ spl66_47 | spl66_137)), inference(backward_demodulation, [], [f1443, f787])).
fof(f2686, plain, (~ spl66_8 | ~ spl66_56), inference(avatar_split_clause, [], [f2656, f823, f619])).
fof(f2656, plain, (~ (e3 = op(e3, e2)) | ~ spl66_56), inference(backward_demodulation, [], [f249, f825])).
fof(f2655, plain, (~ spl66_54 | ~ spl66_58), inference(avatar_split_clause, [], [f2652, f832, f815])).
fof(f2652, plain, (~ (e1 = op(e0, e2)) | ~ spl66_58), inference(backward_demodulation, [], [f262, f834])).
fof(f2644, plain, (~ spl66_54 | ~ spl66_38), inference(avatar_split_clause, [], [f2643, f747, f815])).
fof(f2643, plain, (~ (e1 = op(e0, e2)) | ~ spl66_38), inference(forward_demodulation, [], [f247, f749])).
fof(f2639, plain, (~ spl66_51 | ~ spl66_44 | spl66_147 | ~ spl66_170), inference(avatar_split_clause, [], [f2427, f1605, f1498, f772, f802])).
fof(f2427, plain, (~ (e2 = op(e0, e3)) | (~ spl66_44 | spl66_147 | ~ spl66_170)), inference(backward_demodulation, [], [f2401, f774])).
fof(f2401, plain, (~ (e2 = op(e0, op(e1, e1))) | (spl66_147 | ~ spl66_170)), inference(forward_demodulation, [], [f1500, f1606])).
fof(f1606, plain, ((e0 = op(op(e1, e1), e1)) | ~ spl66_170), inference(avatar_component_clause, [], [f1605])).
fof(f2624, plain, (~ spl66_20 | ~ spl66_4), inference(avatar_split_clause, [], [f2623, f602, f670])).
fof(f2623, plain, (~ (e3 = op(e2, e3)) | ~ spl66_4), inference(forward_demodulation, [], [f258, f604])).
fof(f604, plain, ((e3 = op(e3, e3)) | ~ spl66_4), inference(avatar_component_clause, [], [f602])).
fof(f2620, plain, (~ spl66_19 | ~ spl66_27), inference(avatar_split_clause, [], [f2619, f700, f666])).
fof(f2619, plain, (~ (e2 = op(e2, e3)) | ~ spl66_27), inference(forward_demodulation, [], [f275, f702])).
fof(f2596, plain, (~ spl66_6 | ~ spl66_24 | spl66_154), inference(avatar_split_clause, [], [f2568, f1529, f687, f611])).
fof(f2568, plain, (~ (e1 = op(e3, e2)) | (~ spl66_24 | spl66_154)), inference(backward_demodulation, [], [f1531, f689])).
fof(f1531, plain, (~ (e1 = op(op(e2, e2), e2)) | spl66_154), inference(avatar_component_clause, [], [f1529])).
fof(f2595, plain, (~ spl66_5 | ~ spl66_24 | spl66_164), inference(avatar_split_clause, [], [f2569, f1575, f687, f607])).
fof(f2569, plain, (~ (e0 = op(e3, e2)) | (~ spl66_24 | spl66_164)), inference(backward_demodulation, [], [f1577, f689])).
fof(f2594, plain, (~ spl66_2 | ~ spl66_4 | spl66_158), inference(avatar_split_clause, [], [f2589, f1547, f602, f594])).
fof(f2589, plain, (~ (e1 = op(e3, e3)) | (~ spl66_4 | spl66_158)), inference(backward_demodulation, [], [f1549, f604])).
fof(f2519, plain, (~ spl66_57 | ~ spl66_61), inference(avatar_split_clause, [], [f2509, f845, f828])).
fof(f828, plain, (spl66_57 <=> (e0 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_57])])).
fof(f2509, plain, (~ (e0 = op(e0, e1)) | ~ spl66_61), inference(backward_demodulation, [], [f259, f847])).
fof(f2518, plain, (~ spl66_45 | ~ spl66_61), inference(avatar_split_clause, [], [f2507, f845, f777])).
fof(f2507, plain, (~ (e0 = op(e1, e0)) | ~ spl66_61), inference(backward_demodulation, [], [f235, f847])).
fof(f235, plain, ~ (op(e0, e0) = op(e1, e0)), inference(cnf_transformation, [], [f3])).
fof(f2504, plain, (~ spl66_60 | ~ spl66_44), inference(avatar_split_clause, [], [f2503, f772, f840])).
fof(f2503, plain, (~ (e3 = op(e0, e1)) | ~ spl66_44), inference(forward_demodulation, [], [f241, f774])).
fof(f2475, plain, (~ spl66_2 | ~ spl66_3), inference(avatar_contradiction_clause, [], [f2474])).
fof(f2474, plain, ($false | (~ spl66_2 | ~ spl66_3)), inference(subsumption_resolution, [], [f2473, f286])).
fof(f286, plain, ~ (e1 = e2), inference(cnf_transformation, [], [f4])).
fof(f2473, plain, ((e1 = e2) | (~ spl66_2 | ~ spl66_3)), inference(forward_demodulation, [], [f600, f596])).
fof(f2453, plain, (~ spl66_12 | ~ spl66_28), inference(avatar_split_clause, [], [f2450, f704, f636])).
fof(f2450, plain, (~ (e3 = op(e3, e1)) | ~ spl66_28), inference(backward_demodulation, [], [f246, f706])).
fof(f2429, plain, (~ spl66_12 | ~ spl66_44), inference(avatar_split_clause, [], [f2422, f772, f636])).
fof(f2422, plain, (~ (e3 = op(e3, e1)) | ~ spl66_44), inference(backward_demodulation, [], [f245, f774])).
fof(f2420, plain, (~ spl66_45 | ~ spl66_47), inference(avatar_contradiction_clause, [], [f2419])).
fof(f2419, plain, ($false | (~ spl66_45 | ~ spl66_47)), inference(subsumption_resolution, [], [f2418, f284])).
fof(f284, plain, ~ (e0 = e2), inference(cnf_transformation, [], [f4])).
fof(f2418, plain, ((e0 = e2) | (~ spl66_45 | ~ spl66_47)), inference(backward_demodulation, [], [f787, f779])).
fof(f2405, plain, (~ spl66_57 | ~ spl66_59), inference(avatar_contradiction_clause, [], [f2404])).
fof(f2404, plain, ($false | (~ spl66_57 | ~ spl66_59)), inference(subsumption_resolution, [], [f2403, f284])).
fof(f2403, plain, ((e0 = e2) | (~ spl66_57 | ~ spl66_59)), inference(backward_demodulation, [], [f838, f830])).
fof(f830, plain, ((e0 = op(e0, e1)) | ~ spl66_57), inference(avatar_component_clause, [], [f828])).
fof(f2393, plain, (~ spl66_35 | ~ spl66_2 | spl66_156), inference(avatar_split_clause, [], [f2392, f1538, f594, f734])).
fof(f2392, plain, (~ (e2 = op(e1, e3)) | (~ spl66_2 | spl66_156)), inference(forward_demodulation, [], [f1540, f596])).
fof(f1540, plain, (~ (e2 = op(op(e3, e3), e3)) | spl66_156), inference(avatar_component_clause, [], [f1538])).
fof(f2390, plain, (~ spl66_52 | ~ spl66_64), inference(avatar_split_clause, [], [f2389, f857, f806])).
fof(f2389, plain, (~ (e3 = op(e0, e3)) | ~ spl66_64), inference(forward_demodulation, [], [f261, f859])).
fof(f2364, plain, (~ spl66_2 | spl66_34 | ~ spl66_158), inference(avatar_contradiction_clause, [], [f2363])).
fof(f2363, plain, ($false | (~ spl66_2 | spl66_34 | ~ spl66_158)), inference(subsumption_resolution, [], [f2359, f731])).
fof(f731, plain, (~ (e1 = op(e1, e3)) | spl66_34), inference(avatar_component_clause, [], [f730])).
fof(f2359, plain, ((e1 = op(e1, e3)) | (~ spl66_2 | ~ spl66_158)), inference(backward_demodulation, [], [f1548, f596])).
fof(f1548, plain, ((e1 = op(op(e3, e3), e3)) | ~ spl66_158), inference(avatar_component_clause, [], [f1547])).
fof(f2349, plain, (~ spl66_20 | spl66_129), inference(avatar_contradiction_clause, [], [f2348])).
fof(f2348, plain, ($false | (~ spl66_20 | spl66_129)), inference(subsumption_resolution, [], [f2347, f672])).
fof(f2347, plain, (~ (e3 = op(e2, e3)) | (~ spl66_20 | spl66_129)), inference(backward_demodulation, [], [f1404, f672])).
fof(f1404, plain, (~ (e3 = op(e2, op(e2, e3))) | spl66_129), inference(avatar_component_clause, [], [f1402])).
fof(f1402, plain, (spl66_129 <=> (e3 = op(e2, op(e2, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl66_129])])).
fof(f2337, plain, (~ spl66_29 | spl66_132), inference(avatar_contradiction_clause, [], [f2336])).
fof(f2336, plain, ($false | (~ spl66_29 | spl66_132)), inference(subsumption_resolution, [], [f2331, f711])).
fof(f2331, plain, (~ (e0 = op(e2, e0)) | (~ spl66_29 | spl66_132)), inference(backward_demodulation, [], [f1419, f711])).
fof(f2316, plain, (~ spl66_41 | spl66_57 | ~ spl66_170), inference(avatar_contradiction_clause, [], [f2315])).
fof(f2315, plain, ($false | (~ spl66_41 | spl66_57 | ~ spl66_170)), inference(subsumption_resolution, [], [f2309, f829])).
fof(f829, plain, (~ (e0 = op(e0, e1)) | spl66_57), inference(avatar_component_clause, [], [f828])).
fof(f2309, plain, ((e0 = op(e0, e1)) | (~ spl66_41 | ~ spl66_170)), inference(backward_demodulation, [], [f1606, f762])).
fof(f2311, plain, (~ spl66_25 | ~ spl66_41), inference(avatar_split_clause, [], [f2302, f760, f692])).
fof(f2302, plain, (~ (e0 = op(e2, e1)) | ~ spl66_41), inference(backward_demodulation, [], [f244, f762])).
fof(f2290, plain, (~ spl66_18 | ~ spl66_15 | ~ spl66_64 | spl66_167), inference(avatar_split_clause, [], [f2289, f1590, f857, f649, f662])).
fof(f2289, plain, (~ (e1 = op(e2, e3)) | (~ spl66_15 | ~ spl66_64 | spl66_167)), inference(forward_demodulation, [], [f2282, f651])).
fof(f2282, plain, (~ (e1 = op(op(e3, e0), e3)) | (~ spl66_64 | spl66_167)), inference(backward_demodulation, [], [f1592, f859])).
fof(f2241, plain, (~ spl66_13 | ~ spl66_15), inference(avatar_contradiction_clause, [], [f2240])).
fof(f2240, plain, ($false | (~ spl66_13 | ~ spl66_15)), inference(subsumption_resolution, [], [f2239, f284])).
fof(f2239, plain, ((e0 = e2) | (~ spl66_13 | ~ spl66_15)), inference(forward_demodulation, [], [f651, f643])).
fof(f2238, plain, (~ spl66_25 | ~ spl66_43 | ~ spl66_56 | spl66_149), inference(avatar_contradiction_clause, [], [f2237])).
fof(f2237, plain, ($false | (~ spl66_25 | ~ spl66_43 | ~ spl66_56 | spl66_149)), inference(subsumption_resolution, [], [f2234, f825])).
fof(f2234, plain, (~ (e3 = op(e0, e2)) | (~ spl66_25 | ~ spl66_43 | spl66_149)), inference(backward_demodulation, [], [f2201, f694])).
fof(f2201, plain, (~ (e3 = op(op(e2, e1), e2)) | (~ spl66_43 | spl66_149)), inference(forward_demodulation, [], [f1509, f770])).
fof(f2225, plain, (~ spl66_49 | ~ spl66_51), inference(avatar_contradiction_clause, [], [f2224])).
fof(f2224, plain, ($false | (~ spl66_49 | ~ spl66_51)), inference(subsumption_resolution, [], [f2223, f284])).
fof(f2223, plain, ((e0 = e2) | (~ spl66_49 | ~ spl66_51)), inference(forward_demodulation, [], [f804, f796])).
fof(f2205, plain, (~ spl66_23 | spl66_130), inference(avatar_contradiction_clause, [], [f2204])).
fof(f2204, plain, ($false | (~ spl66_23 | spl66_130)), inference(subsumption_resolution, [], [f2203, f685])).
fof(f2203, plain, (~ (e2 = op(e2, e2)) | (~ spl66_23 | spl66_130)), inference(forward_demodulation, [], [f1409, f685])).
fof(f1409, plain, (~ (e2 = op(e2, op(e2, e2))) | spl66_130), inference(avatar_component_clause, [], [f1407])).
fof(f1407, plain, (spl66_130 <=> (e2 = op(e2, op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl66_130])])).
fof(f2198, plain, (~ spl66_62 | ~ spl66_46), inference(avatar_split_clause, [], [f2197, f781, f849])).
fof(f2197, plain, (~ (op(e0, e0) = e1) | ~ spl66_46), inference(forward_demodulation, [], [f235, f783])).
fof(f2193, plain, (~ spl66_31 | ~ spl66_23), inference(avatar_split_clause, [], [f2122, f683, f717])).
fof(f2122, plain, (~ (e2 = op(e2, e0)) | ~ spl66_23), inference(forward_demodulation, [], [f272, f685])).
fof(f272, plain, ~ (op(e2, e0) = op(e2, e2)), inference(cnf_transformation, [], [f3])).
fof(f2182, plain, (~ spl66_21 | ~ spl66_23), inference(avatar_contradiction_clause, [], [f2181])).
fof(f2181, plain, ($false | (~ spl66_21 | ~ spl66_23)), inference(subsumption_resolution, [], [f2180, f284])).
fof(f2180, plain, ((e0 = e2) | (~ spl66_21 | ~ spl66_23)), inference(backward_demodulation, [], [f685, f677])).
fof(f2177, plain, (~ spl66_41 | ~ spl66_43), inference(avatar_contradiction_clause, [], [f2176])).
fof(f2176, plain, ($false | (~ spl66_41 | ~ spl66_43)), inference(subsumption_resolution, [], [f2175, f284])).
fof(f2175, plain, ((e0 = e2) | (~ spl66_41 | ~ spl66_43)), inference(backward_demodulation, [], [f770, f762])).
fof(f2173, plain, (~ spl66_25 | ~ spl66_43 | spl66_170), inference(avatar_split_clause, [], [f2171, f1605, f768, f692])).
fof(f2171, plain, (~ (e0 = op(e2, e1)) | (~ spl66_43 | spl66_170)), inference(backward_demodulation, [], [f1607, f770])).
fof(f2156, plain, (~ spl66_26 | ~ spl66_58), inference(avatar_split_clause, [], [f2153, f832, f696])).
fof(f2153, plain, (~ (e1 = op(e2, e1)) | ~ spl66_58), inference(backward_demodulation, [], [f242, f834])).
fof(f2128, plain, (~ spl66_37 | ~ spl66_3 | ~ spl66_18 | spl66_165), inference(avatar_split_clause, [], [f2086, f1580, f662, f598, f743])).
fof(f2086, plain, (~ (e0 = op(e1, e2)) | (~ spl66_3 | ~ spl66_18 | spl66_165)), inference(backward_demodulation, [], [f2049, f664])).
fof(f2049, plain, (~ (e0 = op(op(e2, e3), e2)) | (~ spl66_3 | spl66_165)), inference(forward_demodulation, [], [f1582, f600])).
fof(f2118, plain, (~ spl66_47 | ~ spl66_31), inference(avatar_split_clause, [], [f2117, f717, f785])).
fof(f2117, plain, (~ (e2 = op(e1, e0)) | ~ spl66_31), inference(forward_demodulation, [], [f238, f719])).
fof(f2116, plain, (~ spl66_63 | ~ spl66_31), inference(avatar_split_clause, [], [f2115, f717, f853])).
fof(f2115, plain, (~ (op(e0, e0) = e2) | ~ spl66_31), inference(forward_demodulation, [], [f236, f719])).
fof(f2096, plain, (~ spl66_13 | ~ spl66_14), inference(avatar_contradiction_clause, [], [f2095])).
fof(f2095, plain, ($false | (~ spl66_13 | ~ spl66_14)), inference(subsumption_resolution, [], [f2094, f283])).
fof(f2094, plain, ((e0 = e1) | (~ spl66_13 | ~ spl66_14)), inference(backward_demodulation, [], [f647, f643])).
fof(f2080, plain, (~ spl66_7 | ~ spl66_23), inference(avatar_split_clause, [], [f2075, f683, f615])).
fof(f2075, plain, (~ (e2 = op(e3, e2)) | ~ spl66_23), inference(backward_demodulation, [], [f252, f685])).
fof(f2070, plain, (~ spl66_31 | ~ spl66_32), inference(avatar_contradiction_clause, [], [f2069])).
fof(f2069, plain, ($false | (~ spl66_31 | ~ spl66_32)), inference(subsumption_resolution, [], [f2068, f288])).
fof(f288, plain, ~ (e2 = e3), inference(cnf_transformation, [], [f4])).
fof(f2068, plain, ((e2 = e3) | (~ spl66_31 | ~ spl66_32)), inference(backward_demodulation, [], [f723, f719])).
fof(f2067, plain, (~ spl66_37 | ~ spl66_38), inference(avatar_contradiction_clause, [], [f2066])).
fof(f2066, plain, ($false | (~ spl66_37 | ~ spl66_38)), inference(subsumption_resolution, [], [f2065, f283])).
fof(f2065, plain, ((e0 = e1) | (~ spl66_37 | ~ spl66_38)), inference(backward_demodulation, [], [f749, f745])).
fof(f2060, plain, (~ spl66_49 | ~ spl66_50), inference(avatar_contradiction_clause, [], [f2059])).
fof(f2059, plain, ($false | (~ spl66_49 | ~ spl66_50)), inference(subsumption_resolution, [], [f2058, f283])).
fof(f2058, plain, ((e0 = e1) | (~ spl66_49 | ~ spl66_50)), inference(backward_demodulation, [], [f800, f796])).
fof(f2038, plain, (~ spl66_22 | ~ spl66_38), inference(avatar_split_clause, [], [f2037, f747, f679])).
fof(f2037, plain, (~ (e1 = op(e2, e2)) | ~ spl66_38), inference(forward_demodulation, [], [f250, f749])).
fof(f2012, plain, (~ spl66_9 | ~ spl66_12), inference(avatar_contradiction_clause, [], [f2011])).
fof(f2011, plain, ($false | (~ spl66_9 | ~ spl66_12)), inference(subsumption_resolution, [], [f2009, f285])).
fof(f2009, plain, ((e0 = e3) | (~ spl66_9 | ~ spl66_12)), inference(backward_demodulation, [], [f638, f626])).
fof(f1999, plain, (~ spl66_26 | ~ spl66_27), inference(avatar_contradiction_clause, [], [f1998])).
fof(f1998, plain, ($false | (~ spl66_26 | ~ spl66_27)), inference(subsumption_resolution, [], [f1997, f286])).
fof(f1997, plain, ((e1 = e2) | (~ spl66_26 | ~ spl66_27)), inference(backward_demodulation, [], [f702, f698])).
fof(f1990, plain, (~ spl66_38 | ~ spl66_39), inference(avatar_contradiction_clause, [], [f1989])).
fof(f1989, plain, ($false | (~ spl66_38 | ~ spl66_39)), inference(subsumption_resolution, [], [f1988, f286])).
fof(f1988, plain, ((e1 = e2) | (~ spl66_38 | ~ spl66_39)), inference(backward_demodulation, [], [f753, f749])).
fof(f1986, plain, (~ spl66_41 | ~ spl66_42 | spl66_170), inference(avatar_split_clause, [], [f1983, f1605, f764, f760])).
fof(f1983, plain, (~ (e0 = op(e1, e1)) | (~ spl66_42 | spl66_170)), inference(backward_demodulation, [], [f1607, f766])).
fof(f1966, plain, (~ spl66_51 | ~ spl66_3), inference(avatar_split_clause, [], [f1965, f598, f802])).
fof(f1965, plain, (~ (e2 = op(e0, e3)) | ~ spl66_3), inference(forward_demodulation, [], [f255, f600])).
fof(f1959, plain, (~ spl66_43 | ~ spl66_27), inference(avatar_split_clause, [], [f1958, f700, f768])).
fof(f1958, plain, (~ (e2 = op(e1, e1)) | ~ spl66_27), inference(forward_demodulation, [], [f244, f702])).
fof(f1951, plain, (~ spl66_52 | ~ spl66_36), inference(avatar_split_clause, [], [f1949, f738, f806])).
fof(f1949, plain, (~ (e3 = op(e0, e3)) | ~ spl66_36), inference(backward_demodulation, [], [f253, f740])).
fof(f740, plain, ((e3 = op(e1, e3)) | ~ spl66_36), inference(avatar_component_clause, [], [f738])).
fof(f1950, plain, (~ spl66_48 | ~ spl66_36), inference(avatar_split_clause, [], [f1948, f738, f789])).
fof(f1948, plain, (~ (e3 = op(e1, e0)) | ~ spl66_36), inference(backward_demodulation, [], [f267, f740])).
fof(f1937, plain, (~ spl66_51 | ~ spl66_19), inference(avatar_split_clause, [], [f1933, f666, f802])).
fof(f1933, plain, (~ (e2 = op(e0, e3)) | ~ spl66_19), inference(backward_demodulation, [], [f254, f668])).
fof(f668, plain, ((e2 = op(e2, e3)) | ~ spl66_19), inference(avatar_component_clause, [], [f666])).
fof(f1926, plain, (~ spl66_18 | ~ spl66_3 | spl66_158), inference(avatar_split_clause, [], [f1893, f1547, f598, f662])).
fof(f1893, plain, (~ (e1 = op(e2, e3)) | (~ spl66_3 | spl66_158)), inference(backward_demodulation, [], [f1549, f600])).
fof(f1904, plain, (~ spl66_8 | ~ spl66_3 | spl66_143), inference(avatar_split_clause, [], [f1894, f1474, f598, f619])).
fof(f1474, plain, (spl66_143 <=> (e3 = op(e3, op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl66_143])])).
fof(f1894, plain, (~ (e3 = op(e3, e2)) | (~ spl66_3 | spl66_143)), inference(backward_demodulation, [], [f1476, f600])).
fof(f1476, plain, (~ (e3 = op(e3, op(e3, e3))) | spl66_143), inference(avatar_component_clause, [], [f1474])).
fof(f1898, plain, (~ spl66_54 | ~ spl66_3 | ~ spl66_17 | spl66_155), inference(avatar_split_clause, [], [f1897, f1534, f658, f598, f815])).
fof(f1534, plain, (spl66_155 <=> (e1 = op(op(op(e3, e3), e3), op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl66_155])])).
fof(f1897, plain, (~ (e1 = op(e0, e2)) | (~ spl66_3 | ~ spl66_17 | spl66_155)), inference(forward_demodulation, [], [f1892, f660])).
fof(f1892, plain, (~ (e1 = op(op(e2, e3), e2)) | (~ spl66_3 | spl66_155)), inference(backward_demodulation, [], [f1536, f600])).
fof(f1536, plain, (~ (e1 = op(op(op(e3, e3), e3), op(e3, e3))) | spl66_155), inference(avatar_component_clause, [], [f1534])).
fof(f1884, plain, (~ spl66_22 | ~ spl66_23), inference(avatar_contradiction_clause, [], [f1883])).
fof(f1883, plain, ($false | (~ spl66_22 | ~ spl66_23)), inference(subsumption_resolution, [], [f1882, f286])).
fof(f1882, plain, ((e1 = e2) | (~ spl66_22 | ~ spl66_23)), inference(backward_demodulation, [], [f685, f681])).
fof(f1855, plain, (~ spl66_57 | ~ spl66_41), inference(avatar_split_clause, [], [f1854, f760, f828])).
fof(f1854, plain, (~ (e0 = op(e0, e1)) | ~ spl66_41), inference(forward_demodulation, [], [f241, f762])).
fof(f1849, plain, (~ spl66_60 | ~ spl66_41 | spl66_148), inference(avatar_split_clause, [], [f1679, f1502, f760, f840])).
fof(f1679, plain, (~ (e3 = op(e0, e1)) | (~ spl66_41 | spl66_148)), inference(backward_demodulation, [], [f1504, f762])).
fof(f1504, plain, (~ (e3 = op(op(e1, e1), e1)) | spl66_148), inference(avatar_component_clause, [], [f1502])).
fof(f1848, plain, (~ spl66_59 | ~ spl66_41 | spl66_150), inference(avatar_split_clause, [], [f1680, f1511, f760, f836])).
fof(f1680, plain, (~ (e2 = op(e0, e1)) | (~ spl66_41 | spl66_150)), inference(backward_demodulation, [], [f1513, f762])).
fof(f1513, plain, (~ (e2 = op(op(e1, e1), e1)) | spl66_150), inference(avatar_component_clause, [], [f1511])).
fof(f1840, plain, (~ spl66_49 | ~ spl66_17), inference(avatar_split_clause, [], [f1839, f658, f794])).
fof(f1839, plain, (~ (e0 = op(e0, e3)) | ~ spl66_17), inference(forward_demodulation, [], [f254, f660])).
fof(f1834, plain, (~ spl66_45 | ~ spl66_41), inference(avatar_split_clause, [], [f1833, f760, f777])).
fof(f1833, plain, (~ (e0 = op(e1, e0)) | ~ spl66_41), inference(forward_demodulation, [], [f265, f762])).
fof(f1826, plain, (~ spl66_34 | ~ spl66_36), inference(avatar_contradiction_clause, [], [f1825])).
fof(f1825, plain, ($false | (~ spl66_34 | ~ spl66_36)), inference(subsumption_resolution, [], [f1824, f287])).
fof(f1824, plain, ((e1 = e3) | (~ spl66_34 | ~ spl66_36)), inference(forward_demodulation, [], [f740, f732])).
fof(f1823, plain, (~ spl66_31 | ~ spl66_15), inference(avatar_split_clause, [], [f1822, f649, f717])).
fof(f1822, plain, (~ (e2 = op(e2, e0)) | ~ spl66_15), inference(forward_demodulation, [], [f240, f651])).
fof(f1819, plain, (~ spl66_26 | ~ spl66_10), inference(avatar_split_clause, [], [f1818, f628, f696])).
fof(f1818, plain, (~ (e1 = op(e2, e1)) | ~ spl66_10), inference(forward_demodulation, [], [f246, f630])).
fof(f1806, plain, (~ spl66_14 | ~ spl66_10), inference(avatar_split_clause, [], [f1805, f628, f645])).
fof(f1805, plain, (~ (e1 = op(e3, e0)) | ~ spl66_10), inference(forward_demodulation, [], [f277, f630])).
fof(f1797, plain, (~ spl66_10 | ~ spl66_12), inference(avatar_contradiction_clause, [], [f1796])).
fof(f1796, plain, ($false | (~ spl66_10 | ~ spl66_12)), inference(subsumption_resolution, [], [f1795, f287])).
fof(f1795, plain, ((e1 = e3) | (~ spl66_10 | ~ spl66_12)), inference(forward_demodulation, [], [f638, f630])).
fof(f1784, plain, (~ spl66_4 | spl66_143), inference(avatar_contradiction_clause, [], [f1783])).
fof(f1783, plain, ($false | (~ spl66_4 | spl66_143)), inference(subsumption_resolution, [], [f1780, f604])).
fof(f1780, plain, (~ (e3 = op(e3, e3)) | (~ spl66_4 | spl66_143)), inference(backward_demodulation, [], [f1476, f604])).
fof(f1779, plain, (~ spl66_5 | ~ spl66_6), inference(avatar_contradiction_clause, [], [f1778])).
fof(f1778, plain, ($false | (~ spl66_5 | ~ spl66_6)), inference(subsumption_resolution, [], [f1777, f283])).
fof(f1777, plain, ((e0 = e1) | (~ spl66_5 | ~ spl66_6)), inference(backward_demodulation, [], [f613, f609])).
fof(f1776, plain, (~ spl66_6 | ~ spl66_7), inference(avatar_contradiction_clause, [], [f1775])).
fof(f1775, plain, ($false | (~ spl66_6 | ~ spl66_7)), inference(subsumption_resolution, [], [f1774, f286])).
fof(f1774, plain, ((e1 = e2) | (~ spl66_6 | ~ spl66_7)), inference(backward_demodulation, [], [f617, f613])).
fof(f617, plain, ((e2 = op(e3, e2)) | ~ spl66_7), inference(avatar_component_clause, [], [f615])).
fof(f1773, plain, (~ spl66_7 | ~ spl66_8), inference(avatar_contradiction_clause, [], [f1772])).
fof(f1772, plain, ($false | (~ spl66_7 | ~ spl66_8)), inference(subsumption_resolution, [], [f1771, f288])).
fof(f1771, plain, ((e2 = e3) | (~ spl66_7 | ~ spl66_8)), inference(backward_demodulation, [], [f621, f617])).
fof(f1763, plain, (~ spl66_11 | ~ spl66_12), inference(avatar_contradiction_clause, [], [f1762])).
fof(f1762, plain, ($false | (~ spl66_11 | ~ spl66_12)), inference(subsumption_resolution, [], [f1761, f288])).
fof(f1761, plain, ((e2 = e3) | (~ spl66_11 | ~ spl66_12)), inference(backward_demodulation, [], [f638, f634])).
fof(f1751, plain, (~ spl66_15 | ~ spl66_16), inference(avatar_contradiction_clause, [], [f1750])).
fof(f1750, plain, ($false | (~ spl66_15 | ~ spl66_16)), inference(subsumption_resolution, [], [f1749, f288])).
fof(f1749, plain, ((e2 = e3) | (~ spl66_15 | ~ spl66_16)), inference(backward_demodulation, [], [f655, f651])).
fof(f1747, plain, (~ spl66_4 | ~ spl66_16), inference(avatar_split_clause, [], [f1743, f653, f602])).
fof(f1743, plain, (~ (e3 = op(e3, e3)) | ~ spl66_16), inference(backward_demodulation, [], [f279, f655])).
fof(f1737, plain, (~ spl66_18 | ~ spl66_19), inference(avatar_contradiction_clause, [], [f1736])).
fof(f1736, plain, ($false | (~ spl66_18 | ~ spl66_19)), inference(subsumption_resolution, [], [f1735, f286])).
fof(f1735, plain, ((e1 = e2) | (~ spl66_18 | ~ spl66_19)), inference(backward_demodulation, [], [f668, f664])).
fof(f1734, plain, (~ spl66_19 | ~ spl66_20), inference(avatar_contradiction_clause, [], [f1733])).
fof(f1733, plain, ($false | (~ spl66_19 | ~ spl66_20)), inference(subsumption_resolution, [], [f1732, f288])).
fof(f1732, plain, ((e2 = e3) | (~ spl66_19 | ~ spl66_20)), inference(backward_demodulation, [], [f672, f668])).
fof(f1705, plain, (~ spl66_34 | ~ spl66_35), inference(avatar_contradiction_clause, [], [f1704])).
fof(f1704, plain, ($false | (~ spl66_34 | ~ spl66_35)), inference(subsumption_resolution, [], [f1703, f286])).
fof(f1703, plain, ((e1 = e2) | (~ spl66_34 | ~ spl66_35)), inference(backward_demodulation, [], [f736, f732])).
fof(f1702, plain, (~ spl66_35 | ~ spl66_36), inference(avatar_contradiction_clause, [], [f1701])).
fof(f1701, plain, ($false | (~ spl66_35 | ~ spl66_36)), inference(subsumption_resolution, [], [f1700, f288])).
fof(f1700, plain, ((e2 = e3) | (~ spl66_35 | ~ spl66_36)), inference(backward_demodulation, [], [f740, f736])).
fof(f1699, plain, (~ spl66_4 | ~ spl66_36), inference(avatar_split_clause, [], [f1697, f738, f602])).
fof(f1697, plain, (~ (e3 = op(e3, e3)) | ~ spl66_36), inference(backward_demodulation, [], [f257, f740])).
fof(f1674, plain, (~ spl66_33 | ~ spl66_45), inference(avatar_split_clause, [], [f1669, f777, f726])).
fof(f1669, plain, (~ (e0 = op(e1, e3)) | ~ spl66_45), inference(backward_demodulation, [], [f267, f779])).
fof(f1671, plain, (~ spl66_13 | ~ spl66_45), inference(avatar_split_clause, [], [f1666, f777, f641])).
fof(f1666, plain, (~ (e0 = op(e3, e0)) | ~ spl66_45), inference(backward_demodulation, [], [f239, f779])).
fof(f1670, plain, (~ spl66_29 | ~ spl66_45), inference(avatar_split_clause, [], [f1665, f777, f709])).
fof(f1665, plain, (~ (e0 = op(e2, e0)) | ~ spl66_45), inference(backward_demodulation, [], [f238, f779])).
fof(f1664, plain, (~ spl66_50 | ~ spl66_51), inference(avatar_contradiction_clause, [], [f1663])).
fof(f1663, plain, ($false | (~ spl66_50 | ~ spl66_51)), inference(subsumption_resolution, [], [f1662, f286])).
fof(f1662, plain, ((e1 = e2) | (~ spl66_50 | ~ spl66_51)), inference(backward_demodulation, [], [f804, f800])).
fof(f1661, plain, (~ spl66_51 | ~ spl66_52), inference(avatar_contradiction_clause, [], [f1660])).
fof(f1660, plain, ($false | (~ spl66_51 | ~ spl66_52)), inference(subsumption_resolution, [], [f1659, f288])).
fof(f1659, plain, ((e2 = e3) | (~ spl66_51 | ~ spl66_52)), inference(backward_demodulation, [], [f808, f804])).
fof(f1650, plain, (~ spl66_21 | ~ spl66_53), inference(avatar_split_clause, [], [f1646, f811, f675])).
fof(f1646, plain, (~ (e0 = op(e2, e2)) | ~ spl66_53), inference(backward_demodulation, [], [f248, f813])).
fof(f1649, plain, (~ spl66_37 | ~ spl66_53), inference(avatar_split_clause, [], [f1645, f811, f743])).
fof(f1645, plain, (~ (e0 = op(e1, e2)) | ~ spl66_53), inference(backward_demodulation, [], [f247, f813])).
fof(f1644, plain, (~ spl66_49 | ~ spl66_57), inference(avatar_split_clause, [], [f1639, f828, f794])).
fof(f1639, plain, (~ (e0 = op(e0, e3)) | ~ spl66_57), inference(backward_demodulation, [], [f263, f830])).
fof(f1643, plain, (~ spl66_53 | ~ spl66_57), inference(avatar_split_clause, [], [f1638, f828, f811])).
fof(f1638, plain, (~ (e0 = op(e0, e2)) | ~ spl66_57), inference(backward_demodulation, [], [f262, f830])).
fof(f1642, plain, (~ spl66_9 | ~ spl66_57), inference(avatar_split_clause, [], [f1637, f828, f624])).
fof(f1637, plain, (~ (e0 = op(e3, e1)) | ~ spl66_57), inference(backward_demodulation, [], [f243, f830])).
fof(f1641, plain, (~ spl66_25 | ~ spl66_57), inference(avatar_split_clause, [], [f1636, f828, f692])).
fof(f1636, plain, (~ (e0 = op(e2, e1)) | ~ spl66_57), inference(backward_demodulation, [], [f242, f830])).
fof(f1628, plain, (~ spl66_13 | ~ spl66_61), inference(avatar_split_clause, [], [f1619, f845, f641])).
fof(f1619, plain, (~ (e0 = op(e3, e0)) | ~ spl66_61), inference(backward_demodulation, [], [f237, f847])).
fof(f1616, plain, (~ spl66_24 | ~ spl66_151 | ~ spl66_164), inference(avatar_split_clause, [], [f588, f1575, f1516, f687])).
fof(f588, plain, (~ (e0 = op(op(e2, e2), e2)) | ~ (e1 = op(op(op(e2, e2), e2), op(e2, e2))) | ~ (e3 = op(e2, e2))), inference(cnf_transformation, [], [f53])).
fof(f53, plain, (~ (e0 = op(op(e2, e2), e2)) | ~ (e1 = op(op(op(e2, e2), e2), op(e2, e2))) | ~ (e3 = op(e2, e2))), inference(ennf_transformation, [], [f29])).
fof(f29, plain, ~ ((e0 = op(op(e2, e2), e2)) & (e1 = op(op(op(e2, e2), e2), op(e2, e2))) & (e3 = op(e2, e2))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG143+1.p', ax29)).
fof(f1615, plain, (~ spl66_24 | ~ spl66_163 | ~ spl66_154), inference(avatar_split_clause, [], [f587, f1529, f1570, f687])).
fof(f587, plain, (~ (e1 = op(op(e2, e2), e2)) | ~ (e0 = op(op(op(e2, e2), e2), op(e2, e2))) | ~ (e3 = op(e2, e2))), inference(cnf_transformation, [], [f52])).
fof(f52, plain, (~ (e1 = op(op(e2, e2), e2)) | ~ (e0 = op(op(op(e2, e2), e2), op(e2, e2))) | ~ (e3 = op(e2, e2))), inference(ennf_transformation, [], [f28])).
fof(f28, plain, ~ ((e1 = op(op(e2, e2), e2)) & (e0 = op(op(op(e2, e2), e2), op(e2, e2))) & (e3 = op(e2, e2))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG143+1.p', ax28)).
fof(f1614, plain, (~ spl66_44 | ~ spl66_147 | ~ spl66_170), inference(avatar_split_clause, [], [f586, f1605, f1498, f772])).
fof(f586, plain, (~ (e0 = op(op(e1, e1), e1)) | ~ (e2 = op(op(op(e1, e1), e1), op(e1, e1))) | ~ (e3 = op(e1, e1))), inference(cnf_transformation, [], [f51])).
fof(f51, plain, (~ (e0 = op(op(e1, e1), e1)) | ~ (e2 = op(op(op(e1, e1), e1), op(e1, e1))) | ~ (e3 = op(e1, e1))), inference(ennf_transformation, [], [f27])).
fof(f27, plain, ~ ((e0 = op(op(e1, e1), e1)) & (e2 = op(op(op(e1, e1), e1), op(e1, e1))) & (e3 = op(e1, e1))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG143+1.p', ax27)).
fof(f1611, plain, (~ spl66_64 | ~ spl66_167 | ~ spl66_162), inference(avatar_split_clause, [], [f583, f1565, f1590, f857])).
fof(f583, plain, (~ (e2 = op(op(e0, e0), e0)) | ~ (e1 = op(op(op(e0, e0), e0), op(e0, e0))) | ~ (op(e0, e0) = e3)), inference(cnf_transformation, [], [f48])).
fof(f48, plain, (~ (e2 = op(op(e0, e0), e0)) | ~ (e1 = op(op(op(e0, e0), e0), op(e0, e0))) | ~ (op(e0, e0) = e3)), inference(ennf_transformation, [], [f24])).
fof(f24, plain, ~ ((e2 = op(op(e0, e0), e0)) & (e1 = op(op(op(e0, e0), e0), op(e0, e0))) & (op(e0, e0) = e3)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG143+1.p', ax24)).
fof(f1610, plain, (~ spl66_3 | ~ spl66_155 | ~ spl66_166), inference(avatar_split_clause, [], [f582, f1585, f1534, f598])).
fof(f582, plain, (~ (e0 = op(op(e3, e3), e3)) | ~ (e1 = op(op(op(e3, e3), e3), op(e3, e3))) | ~ (e2 = op(e3, e3))), inference(cnf_transformation, [], [f47])).
fof(f47, plain, (~ (e0 = op(op(e3, e3), e3)) | ~ (e1 = op(op(op(e3, e3), e3), op(e3, e3))) | ~ (e2 = op(e3, e3))), inference(ennf_transformation, [], [f23])).
fof(f23, plain, ~ ((e0 = op(op(e3, e3), e3)) & (e1 = op(op(op(e3, e3), e3), op(e3, e3))) & (e2 = op(e3, e3))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG143+1.p', ax23)).
fof(f1609, plain, (~ spl66_3 | ~ spl66_165 | ~ spl66_158), inference(avatar_split_clause, [], [f581, f1547, f1580, f598])).
fof(f581, plain, (~ (e1 = op(op(e3, e3), e3)) | ~ (e0 = op(op(op(e3, e3), e3), op(e3, e3))) | ~ (e2 = op(e3, e3))), inference(cnf_transformation, [], [f46])).
fof(f46, plain, (~ (e1 = op(op(e3, e3), e3)) | ~ (e0 = op(op(op(e3, e3), e3), op(e3, e3))) | ~ (e2 = op(e3, e3))), inference(ennf_transformation, [], [f22])).
fof(f22, plain, ~ ((e1 = op(op(e3, e3), e3)) & (e0 = op(op(op(e3, e3), e3), op(e3, e3))) & (e2 = op(e3, e3))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG143+1.p', ax22)).
fof(f1608, plain, (~ spl66_43 | ~ spl66_149 | ~ spl66_170), inference(avatar_split_clause, [], [f580, f1605, f1507, f768])).
fof(f580, plain, (~ (e0 = op(op(e1, e1), e1)) | ~ (e3 = op(op(op(e1, e1), e1), op(e1, e1))) | ~ (e2 = op(e1, e1))), inference(cnf_transformation, [], [f45])).
fof(f45, plain, (~ (e0 = op(op(e1, e1), e1)) | ~ (e3 = op(op(op(e1, e1), e1), op(e1, e1))) | ~ (e2 = op(e1, e1))), inference(ennf_transformation, [], [f21])).
fof(f21, plain, ~ ((e0 = op(op(e1, e1), e1)) & (e3 = op(op(op(e1, e1), e1), op(e1, e1))) & (e2 = op(e1, e1))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG143+1.p', ax21)).
fof(f1593, plain, (~ spl66_63 | ~ spl66_167 | ~ spl66_160), inference(avatar_split_clause, [], [f577, f1556, f1590, f853])).
fof(f577, plain, (~ (e3 = op(op(e0, e0), e0)) | ~ (e1 = op(op(op(e0, e0), e0), op(e0, e0))) | ~ (op(e0, e0) = e2)), inference(cnf_transformation, [], [f42])).
fof(f42, plain, (~ (e3 = op(op(e0, e0), e0)) | ~ (e1 = op(op(op(e0, e0), e0), op(e0, e0))) | ~ (op(e0, e0) = e2)), inference(ennf_transformation, [], [f18])).
fof(f18, plain, ~ ((e3 = op(op(e0, e0), e0)) & (e1 = op(op(op(e0, e0), e0), op(e0, e0))) & (op(e0, e0) = e2)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG143+1.p', ax18)).
fof(f1588, plain, (~ spl66_2 | ~ spl66_157 | ~ spl66_166), inference(avatar_split_clause, [], [f576, f1585, f1543, f594])).
fof(f576, plain, (~ (e0 = op(op(e3, e3), e3)) | ~ (e2 = op(op(op(e3, e3), e3), op(e3, e3))) | ~ (e1 = op(e3, e3))), inference(cnf_transformation, [], [f41])).
fof(f41, plain, (~ (e0 = op(op(e3, e3), e3)) | ~ (e2 = op(op(op(e3, e3), e3), op(e3, e3))) | ~ (e1 = op(e3, e3))), inference(ennf_transformation, [], [f17])).
fof(f17, plain, ~ ((e0 = op(op(e3, e3), e3)) & (e2 = op(op(op(e3, e3), e3), op(e3, e3))) & (e1 = op(e3, e3))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG143+1.p', ax17)).
fof(f1583, plain, (~ spl66_2 | ~ spl66_165 | ~ spl66_156), inference(avatar_split_clause, [], [f575, f1538, f1580, f594])).
fof(f575, plain, (~ (e2 = op(op(e3, e3), e3)) | ~ (e0 = op(op(op(e3, e3), e3), op(e3, e3))) | ~ (e1 = op(e3, e3))), inference(cnf_transformation, [], [f40])).
fof(f40, plain, (~ (e2 = op(op(e3, e3), e3)) | ~ (e0 = op(op(op(e3, e3), e3), op(e3, e3))) | ~ (e1 = op(e3, e3))), inference(ennf_transformation, [], [f16])).
fof(f16, plain, ~ ((e2 = op(op(e3, e3), e3)) & (e0 = op(op(op(e3, e3), e3), op(e3, e3))) & (e1 = op(e3, e3))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG143+1.p', ax16)).
fof(f1578, plain, (~ spl66_22 | ~ spl66_153 | ~ spl66_164), inference(avatar_split_clause, [], [f574, f1575, f1525, f679])).
fof(f574, plain, (~ (e0 = op(op(e2, e2), e2)) | ~ (e3 = op(op(op(e2, e2), e2), op(e2, e2))) | ~ (e1 = op(e2, e2))), inference(cnf_transformation, [], [f39])).
fof(f39, plain, (~ (e0 = op(op(e2, e2), e2)) | ~ (e3 = op(op(op(e2, e2), e2), op(e2, e2))) | ~ (e1 = op(e2, e2))), inference(ennf_transformation, [], [f15])).
fof(f15, plain, ~ ((e0 = op(op(e2, e2), e2)) & (e3 = op(op(op(e2, e2), e2), op(e2, e2))) & (e1 = op(e2, e2))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG143+1.p', ax15)).
fof(f1573, plain, (~ spl66_22 | ~ spl66_163 | ~ spl66_152), inference(avatar_split_clause, [], [f573, f1520, f1570, f679])).
fof(f573, plain, (~ (e3 = op(op(e2, e2), e2)) | ~ (e0 = op(op(op(e2, e2), e2), op(e2, e2))) | ~ (e1 = op(e2, e2))), inference(cnf_transformation, [], [f38])).
fof(f38, plain, (~ (e3 = op(op(e2, e2), e2)) | ~ (e0 = op(op(op(e2, e2), e2), op(e2, e2))) | ~ (e1 = op(e2, e2))), inference(ennf_transformation, [], [f14])).
fof(f14, plain, ~ ((e3 = op(op(e2, e2), e2)) & (e0 = op(op(op(e2, e2), e2), op(e2, e2))) & (e1 = op(e2, e2))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG143+1.p', ax14)).
fof(f1514, plain, (~ spl66_41 | ~ spl66_149 | ~ spl66_150), inference(avatar_split_clause, [], [f566, f1511, f1507, f760])).
fof(f566, plain, (~ (e2 = op(op(e1, e1), e1)) | ~ (e3 = op(op(op(e1, e1), e1), op(e1, e1))) | ~ (e0 = op(e1, e1))), inference(cnf_transformation, [], [f31])).
fof(f31, plain, (~ (e2 = op(op(e1, e1), e1)) | ~ (e3 = op(op(op(e1, e1), e1), op(e1, e1))) | ~ (e0 = op(e1, e1))), inference(ennf_transformation, [], [f7])).
fof(f7, plain, ~ ((e2 = op(op(e1, e1), e1)) & (e3 = op(op(op(e1, e1), e1), op(e1, e1))) & (e0 = op(e1, e1))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG143+1.p', ax7)).
fof(f1505, plain, (~ spl66_41 | ~ spl66_147 | ~ spl66_148), inference(avatar_split_clause, [], [f565, f1502, f1498, f760])).
fof(f565, plain, (~ (e3 = op(op(e1, e1), e1)) | ~ (e2 = op(op(op(e1, e1), e1), op(e1, e1))) | ~ (e0 = op(e1, e1))), inference(cnf_transformation, [], [f30])).
fof(f30, plain, (~ (e3 = op(op(e1, e1), e1)) | ~ (e2 = op(op(op(e1, e1), e1), op(e1, e1))) | ~ (e0 = op(e1, e1))), inference(ennf_transformation, [], [f6])).
fof(f6, plain, ~ ((e3 = op(op(e1, e1), e1)) & (e2 = op(op(op(e1, e1), e1), op(e1, e1))) & (e0 = op(e1, e1))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG143+1.p', ax6)).
fof(f1496, plain, (spl66_61 | spl66_46 | spl66_31 | spl66_16), inference(avatar_split_clause, [], [f553, f653, f717, f781, f845])).
fof(f553, plain, ((e3 = op(e3, e0)) | (e2 = op(e2, e0)) | (e1 = op(e1, e0)) | (e0 = op(e0, e0))), inference(cnf_transformation, [], [f120])).
fof(f120, plain, ((((e3 = op(e3, e3)) & (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3) & ((~ (e3 = op(e3, op(e3, e3))) & ~ (e2 = op(e3, op(e3, e2))) & ~ (e1 = op(e3, op(e3, e1))) & ~ (e0 = op(e3, op(e3, e0)))) | sP2 | sP1 | sP0) & ((e3 = op(e3, e3)) | (e2 = op(e2, e3)) | (e1 = op(e1, e3)) | (e0 = op(e0, e3))) & ((e3 = op(e3, e2)) | (e2 = op(e2, e2)) | (e1 = op(e1, e2)) | (e0 = op(e0, e2))) & ((e3 = op(e3, e1)) | (e2 = op(e2, e1)) | (e1 = op(e1, e1)) | (e0 = op(e0, e1))) & ((e3 = op(e3, e0)) | (e2 = op(e2, e0)) | (e1 = op(e1, e0)) | (e0 = op(e0, e0)))), inference(definition_folding, [], [f5, e119, e118, e117, e116, e115, e114, e113, e112, e111, e110, e109, e108, e107, e106, e105, e104, e103, e102, e101, e100, e99, e98, e97, e96, e95, e94, e93, e92, e91, e90, e89, e88, e87, e86, e85, e84, e83, e82, e81, e80, e79, e78, e77, e76, e75, e74, e73, e72, e71, e70, e69, e68, e67, e66, e65, e64, e63, e62, e61, e60, e59, e58, e57, e56, e55, e54])).
fof(f54, plain, ((~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0)))) | ~ sP0), inference(usedef, [], [e54])).
fof(e54, plain, (sP0 <=> (~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f55, plain, ((~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0)))) | ~ sP1), inference(usedef, [], [e55])).
fof(e55, plain, (sP1 <=> (~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f56, plain, ((~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0)))) | ~ sP2), inference(usedef, [], [e56])).
fof(e56, plain, (sP2 <=> (~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f57, plain, (((e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP3), inference(usedef, [], [e57])).
fof(e57, plain, (sP3 <=> ((e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f58, plain, (((e0 = op(e0, e1)) & (e0 = op(e1, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP4), inference(usedef, [], [e58])).
fof(e58, plain, (sP4 <=> ((e0 = op(e0, e1)) & (e0 = op(e1, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f59, plain, (((e0 = op(e0, e2)) & (e0 = op(e2, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP5), inference(usedef, [], [e59])).
fof(e59, plain, (sP5 <=> ((e0 = op(e0, e2)) & (e0 = op(e2, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f60, plain, (((e0 = op(e0, e3)) & (e0 = op(e3, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP6), inference(usedef, [], [e60])).
fof(e60, plain, (sP6 <=> ((e0 = op(e0, e3)) & (e0 = op(e3, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f61, plain, (((e1 = op(e1, e0)) & (e1 = op(e0, e1)) & (e0 = op(e1, e0)) & (e0 = op(e0, e1))) | ~ sP7), inference(usedef, [], [e61])).
fof(e61, plain, (sP7 <=> ((e1 = op(e1, e0)) & (e1 = op(e0, e1)) & (e0 = op(e1, e0)) & (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f62, plain, (((e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e0 = op(e1, e0)) & (e0 = op(e0, e1))) | ~ sP8), inference(usedef, [], [e62])).
fof(e62, plain, (sP8 <=> ((e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e0 = op(e1, e0)) & (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f63, plain, (((e1 = op(e1, e2)) & (e1 = op(e2, e1)) & (e0 = op(e1, e0)) & (e0 = op(e0, e1))) | ~ sP9), inference(usedef, [], [e63])).
fof(e63, plain, (sP9 <=> ((e1 = op(e1, e2)) & (e1 = op(e2, e1)) & (e0 = op(e1, e0)) & (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f64, plain, (((e1 = op(e1, e3)) & (e1 = op(e3, e1)) & (e0 = op(e1, e0)) & (e0 = op(e0, e1))) | ~ sP10), inference(usedef, [], [e64])).
fof(e64, plain, (sP10 <=> ((e1 = op(e1, e3)) & (e1 = op(e3, e1)) & (e0 = op(e1, e0)) & (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f65, plain, (((e2 = op(e2, e0)) & (e2 = op(e0, e2)) & (e0 = op(e2, e0)) & (e0 = op(e0, e2))) | ~ sP11), inference(usedef, [], [e65])).
fof(e65, plain, (sP11 <=> ((e2 = op(e2, e0)) & (e2 = op(e0, e2)) & (e0 = op(e2, e0)) & (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f66, plain, (((e2 = op(e2, e1)) & (e2 = op(e1, e2)) & (e0 = op(e2, e0)) & (e0 = op(e0, e2))) | ~ sP12), inference(usedef, [], [e66])).
fof(e66, plain, (sP12 <=> ((e2 = op(e2, e1)) & (e2 = op(e1, e2)) & (e0 = op(e2, e0)) & (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f67, plain, (((e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e0 = op(e2, e0)) & (e0 = op(e0, e2))) | ~ sP13), inference(usedef, [], [e67])).
fof(e67, plain, (sP13 <=> ((e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e0 = op(e2, e0)) & (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f68, plain, (((e2 = op(e2, e3)) & (e2 = op(e3, e2)) & (e0 = op(e2, e0)) & (e0 = op(e0, e2))) | ~ sP14), inference(usedef, [], [e68])).
fof(e68, plain, (sP14 <=> ((e2 = op(e2, e3)) & (e2 = op(e3, e2)) & (e0 = op(e2, e0)) & (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f69, plain, (((e3 = op(e3, e0)) & (e3 = op(e0, e3)) & (e0 = op(e3, e0)) & (e0 = op(e0, e3))) | ~ sP15), inference(usedef, [], [e69])).
fof(e69, plain, (sP15 <=> ((e3 = op(e3, e0)) & (e3 = op(e0, e3)) & (e0 = op(e3, e0)) & (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f70, plain, (((e3 = op(e3, e1)) & (e3 = op(e1, e3)) & (e0 = op(e3, e0)) & (e0 = op(e0, e3))) | ~ sP16), inference(usedef, [], [e70])).
fof(e70, plain, (sP16 <=> ((e3 = op(e3, e1)) & (e3 = op(e1, e3)) & (e0 = op(e3, e0)) & (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f71, plain, (((e3 = op(e3, e2)) & (e3 = op(e2, e3)) & (e0 = op(e3, e0)) & (e0 = op(e0, e3))) | ~ sP17), inference(usedef, [], [e71])).
fof(e71, plain, (sP17 <=> ((e3 = op(e3, e2)) & (e3 = op(e2, e3)) & (e0 = op(e3, e0)) & (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f72, plain, (((e3 = op(e3, e3)) & (e3 = op(e3, e3)) & (e0 = op(e3, e0)) & (e0 = op(e0, e3))) | ~ sP18), inference(usedef, [], [e72])).
fof(e72, plain, (sP18 <=> ((e3 = op(e3, e3)) & (e3 = op(e3, e3)) & (e0 = op(e3, e0)) & (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f73, plain, (((e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e1 = op(e0, e1)) & (e1 = op(e1, e0))) | ~ sP19), inference(usedef, [], [e73])).
fof(e73, plain, (sP19 <=> ((e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e1 = op(e0, e1)) & (e1 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP19])])).
fof(f74, plain, (((e0 = op(e0, e1)) & (e0 = op(e1, e0)) & (e1 = op(e0, e1)) & (e1 = op(e1, e0))) | ~ sP20), inference(usedef, [], [e74])).
fof(e74, plain, (sP20 <=> ((e0 = op(e0, e1)) & (e0 = op(e1, e0)) & (e1 = op(e0, e1)) & (e1 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP20])])).
fof(f75, plain, (((e0 = op(e0, e2)) & (e0 = op(e2, e0)) & (e1 = op(e0, e1)) & (e1 = op(e1, e0))) | ~ sP21), inference(usedef, [], [e75])).
fof(e75, plain, (sP21 <=> ((e0 = op(e0, e2)) & (e0 = op(e2, e0)) & (e1 = op(e0, e1)) & (e1 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP21])])).
fof(f76, plain, (((e0 = op(e0, e3)) & (e0 = op(e3, e0)) & (e1 = op(e0, e1)) & (e1 = op(e1, e0))) | ~ sP22), inference(usedef, [], [e76])).
fof(e76, plain, (sP22 <=> ((e0 = op(e0, e3)) & (e0 = op(e3, e0)) & (e1 = op(e0, e1)) & (e1 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP22])])).
fof(f77, plain, (((e1 = op(e1, e0)) & (e1 = op(e0, e1)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP23), inference(usedef, [], [e77])).
fof(e77, plain, (sP23 <=> ((e1 = op(e1, e0)) & (e1 = op(e0, e1)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP23])])).
fof(f78, plain, (((e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP24), inference(usedef, [], [e78])).
fof(e78, plain, (sP24 <=> ((e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP24])])).
fof(f79, plain, (((e1 = op(e1, e2)) & (e1 = op(e2, e1)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP25), inference(usedef, [], [e79])).
fof(e79, plain, (sP25 <=> ((e1 = op(e1, e2)) & (e1 = op(e2, e1)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP25])])).
fof(f80, plain, (((e1 = op(e1, e3)) & (e1 = op(e3, e1)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP26), inference(usedef, [], [e80])).
fof(e80, plain, (sP26 <=> ((e1 = op(e1, e3)) & (e1 = op(e3, e1)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP26])])).
fof(f81, plain, (((e2 = op(e2, e0)) & (e2 = op(e0, e2)) & (e1 = op(e2, e1)) & (e1 = op(e1, e2))) | ~ sP27), inference(usedef, [], [e81])).
fof(e81, plain, (sP27 <=> ((e2 = op(e2, e0)) & (e2 = op(e0, e2)) & (e1 = op(e2, e1)) & (e1 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP27])])).
fof(f82, plain, (((e2 = op(e2, e1)) & (e2 = op(e1, e2)) & (e1 = op(e2, e1)) & (e1 = op(e1, e2))) | ~ sP28), inference(usedef, [], [e82])).
fof(e82, plain, (sP28 <=> ((e2 = op(e2, e1)) & (e2 = op(e1, e2)) & (e1 = op(e2, e1)) & (e1 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP28])])).
fof(f83, plain, (((e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e1 = op(e2, e1)) & (e1 = op(e1, e2))) | ~ sP29), inference(usedef, [], [e83])).
fof(e83, plain, (sP29 <=> ((e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e1 = op(e2, e1)) & (e1 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP29])])).
fof(f84, plain, (((e2 = op(e2, e3)) & (e2 = op(e3, e2)) & (e1 = op(e2, e1)) & (e1 = op(e1, e2))) | ~ sP30), inference(usedef, [], [e84])).
fof(e84, plain, (sP30 <=> ((e2 = op(e2, e3)) & (e2 = op(e3, e2)) & (e1 = op(e2, e1)) & (e1 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP30])])).
fof(f85, plain, (((e3 = op(e3, e0)) & (e3 = op(e0, e3)) & (e1 = op(e3, e1)) & (e1 = op(e1, e3))) | ~ sP31), inference(usedef, [], [e85])).
fof(e85, plain, (sP31 <=> ((e3 = op(e3, e0)) & (e3 = op(e0, e3)) & (e1 = op(e3, e1)) & (e1 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP31])])).
fof(f86, plain, (((e3 = op(e3, e1)) & (e3 = op(e1, e3)) & (e1 = op(e3, e1)) & (e1 = op(e1, e3))) | ~ sP32), inference(usedef, [], [e86])).
fof(e86, plain, (sP32 <=> ((e3 = op(e3, e1)) & (e3 = op(e1, e3)) & (e1 = op(e3, e1)) & (e1 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP32])])).
fof(f87, plain, (((e3 = op(e3, e2)) & (e3 = op(e2, e3)) & (e1 = op(e3, e1)) & (e1 = op(e1, e3))) | ~ sP33), inference(usedef, [], [e87])).
fof(e87, plain, (sP33 <=> ((e3 = op(e3, e2)) & (e3 = op(e2, e3)) & (e1 = op(e3, e1)) & (e1 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP33])])).
fof(f88, plain, (((e3 = op(e3, e3)) & (e3 = op(e3, e3)) & (e1 = op(e3, e1)) & (e1 = op(e1, e3))) | ~ sP34), inference(usedef, [], [e88])).
fof(e88, plain, (sP34 <=> ((e3 = op(e3, e3)) & (e3 = op(e3, e3)) & (e1 = op(e3, e1)) & (e1 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP34])])).
fof(f89, plain, (((e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e2 = op(e0, e2)) & (e2 = op(e2, e0))) | ~ sP35), inference(usedef, [], [e89])).
fof(e89, plain, (sP35 <=> ((e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e2 = op(e0, e2)) & (e2 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP35])])).
fof(f90, plain, (((e0 = op(e0, e1)) & (e0 = op(e1, e0)) & (e2 = op(e0, e2)) & (e2 = op(e2, e0))) | ~ sP36), inference(usedef, [], [e90])).
fof(e90, plain, (sP36 <=> ((e0 = op(e0, e1)) & (e0 = op(e1, e0)) & (e2 = op(e0, e2)) & (e2 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP36])])).
fof(f91, plain, (((e0 = op(e0, e2)) & (e0 = op(e2, e0)) & (e2 = op(e0, e2)) & (e2 = op(e2, e0))) | ~ sP37), inference(usedef, [], [e91])).
fof(e91, plain, (sP37 <=> ((e0 = op(e0, e2)) & (e0 = op(e2, e0)) & (e2 = op(e0, e2)) & (e2 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP37])])).
fof(f92, plain, (((e0 = op(e0, e3)) & (e0 = op(e3, e0)) & (e2 = op(e0, e2)) & (e2 = op(e2, e0))) | ~ sP38), inference(usedef, [], [e92])).
fof(e92, plain, (sP38 <=> ((e0 = op(e0, e3)) & (e0 = op(e3, e0)) & (e2 = op(e0, e2)) & (e2 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP38])])).
fof(f93, plain, (((e1 = op(e1, e0)) & (e1 = op(e0, e1)) & (e2 = op(e1, e2)) & (e2 = op(e2, e1))) | ~ sP39), inference(usedef, [], [e93])).
fof(e93, plain, (sP39 <=> ((e1 = op(e1, e0)) & (e1 = op(e0, e1)) & (e2 = op(e1, e2)) & (e2 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP39])])).
fof(f94, plain, (((e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e2 = op(e1, e2)) & (e2 = op(e2, e1))) | ~ sP40), inference(usedef, [], [e94])).
fof(e94, plain, (sP40 <=> ((e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e2 = op(e1, e2)) & (e2 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP40])])).
fof(f95, plain, (((e1 = op(e1, e2)) & (e1 = op(e2, e1)) & (e2 = op(e1, e2)) & (e2 = op(e2, e1))) | ~ sP41), inference(usedef, [], [e95])).
fof(e95, plain, (sP41 <=> ((e1 = op(e1, e2)) & (e1 = op(e2, e1)) & (e2 = op(e1, e2)) & (e2 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP41])])).
fof(f96, plain, (((e1 = op(e1, e3)) & (e1 = op(e3, e1)) & (e2 = op(e1, e2)) & (e2 = op(e2, e1))) | ~ sP42), inference(usedef, [], [e96])).
fof(e96, plain, (sP42 <=> ((e1 = op(e1, e3)) & (e1 = op(e3, e1)) & (e2 = op(e1, e2)) & (e2 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP42])])).
fof(f97, plain, (((e2 = op(e2, e0)) & (e2 = op(e0, e2)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP43), inference(usedef, [], [e97])).
fof(e97, plain, (sP43 <=> ((e2 = op(e2, e0)) & (e2 = op(e0, e2)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP43])])).
fof(f98, plain, (((e2 = op(e2, e1)) & (e2 = op(e1, e2)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP44), inference(usedef, [], [e98])).
fof(e98, plain, (sP44 <=> ((e2 = op(e2, e1)) & (e2 = op(e1, e2)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP44])])).
fof(f99, plain, (((e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP45), inference(usedef, [], [e99])).
fof(e99, plain, (sP45 <=> ((e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP45])])).
fof(f100, plain, (((e2 = op(e2, e3)) & (e2 = op(e3, e2)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP46), inference(usedef, [], [e100])).
fof(e100, plain, (sP46 <=> ((e2 = op(e2, e3)) & (e2 = op(e3, e2)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP46])])).
fof(f101, plain, (((e3 = op(e3, e0)) & (e3 = op(e0, e3)) & (e2 = op(e3, e2)) & (e2 = op(e2, e3))) | ~ sP47), inference(usedef, [], [e101])).
fof(e101, plain, (sP47 <=> ((e3 = op(e3, e0)) & (e3 = op(e0, e3)) & (e2 = op(e3, e2)) & (e2 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP47])])).
fof(f102, plain, (((e3 = op(e3, e1)) & (e3 = op(e1, e3)) & (e2 = op(e3, e2)) & (e2 = op(e2, e3))) | ~ sP48), inference(usedef, [], [e102])).
fof(e102, plain, (sP48 <=> ((e3 = op(e3, e1)) & (e3 = op(e1, e3)) & (e2 = op(e3, e2)) & (e2 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP48])])).
fof(f103, plain, (((e3 = op(e3, e2)) & (e3 = op(e2, e3)) & (e2 = op(e3, e2)) & (e2 = op(e2, e3))) | ~ sP49), inference(usedef, [], [e103])).
fof(e103, plain, (sP49 <=> ((e3 = op(e3, e2)) & (e3 = op(e2, e3)) & (e2 = op(e3, e2)) & (e2 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP49])])).
fof(f104, plain, (((e3 = op(e3, e3)) & (e3 = op(e3, e3)) & (e2 = op(e3, e2)) & (e2 = op(e2, e3))) | ~ sP50), inference(usedef, [], [e104])).
fof(e104, plain, (sP50 <=> ((e3 = op(e3, e3)) & (e3 = op(e3, e3)) & (e2 = op(e3, e2)) & (e2 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP50])])).
fof(f105, plain, (((e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e3 = op(e0, e3)) & (e3 = op(e3, e0))) | ~ sP51), inference(usedef, [], [e105])).
fof(e105, plain, (sP51 <=> ((e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e3 = op(e0, e3)) & (e3 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP51])])).
fof(f106, plain, (((e0 = op(e0, e1)) & (e0 = op(e1, e0)) & (e3 = op(e0, e3)) & (e3 = op(e3, e0))) | ~ sP52), inference(usedef, [], [e106])).
fof(e106, plain, (sP52 <=> ((e0 = op(e0, e1)) & (e0 = op(e1, e0)) & (e3 = op(e0, e3)) & (e3 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP52])])).
fof(f107, plain, (((e0 = op(e0, e2)) & (e0 = op(e2, e0)) & (e3 = op(e0, e3)) & (e3 = op(e3, e0))) | ~ sP53), inference(usedef, [], [e107])).
fof(e107, plain, (sP53 <=> ((e0 = op(e0, e2)) & (e0 = op(e2, e0)) & (e3 = op(e0, e3)) & (e3 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP53])])).
fof(f108, plain, (((e0 = op(e0, e3)) & (e0 = op(e3, e0)) & (e3 = op(e0, e3)) & (e3 = op(e3, e0))) | ~ sP54), inference(usedef, [], [e108])).
fof(e108, plain, (sP54 <=> ((e0 = op(e0, e3)) & (e0 = op(e3, e0)) & (e3 = op(e0, e3)) & (e3 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP54])])).
fof(f109, plain, (((e1 = op(e1, e0)) & (e1 = op(e0, e1)) & (e3 = op(e1, e3)) & (e3 = op(e3, e1))) | ~ sP55), inference(usedef, [], [e109])).
fof(e109, plain, (sP55 <=> ((e1 = op(e1, e0)) & (e1 = op(e0, e1)) & (e3 = op(e1, e3)) & (e3 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP55])])).
fof(f110, plain, (((e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e3 = op(e1, e3)) & (e3 = op(e3, e1))) | ~ sP56), inference(usedef, [], [e110])).
fof(e110, plain, (sP56 <=> ((e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e3 = op(e1, e3)) & (e3 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP56])])).
fof(f111, plain, (((e1 = op(e1, e2)) & (e1 = op(e2, e1)) & (e3 = op(e1, e3)) & (e3 = op(e3, e1))) | ~ sP57), inference(usedef, [], [e111])).
fof(e111, plain, (sP57 <=> ((e1 = op(e1, e2)) & (e1 = op(e2, e1)) & (e3 = op(e1, e3)) & (e3 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP57])])).
fof(f112, plain, (((e1 = op(e1, e3)) & (e1 = op(e3, e1)) & (e3 = op(e1, e3)) & (e3 = op(e3, e1))) | ~ sP58), inference(usedef, [], [e112])).
fof(e112, plain, (sP58 <=> ((e1 = op(e1, e3)) & (e1 = op(e3, e1)) & (e3 = op(e1, e3)) & (e3 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP58])])).
fof(f113, plain, (((e2 = op(e2, e0)) & (e2 = op(e0, e2)) & (e3 = op(e2, e3)) & (e3 = op(e3, e2))) | ~ sP59), inference(usedef, [], [e113])).
fof(e113, plain, (sP59 <=> ((e2 = op(e2, e0)) & (e2 = op(e0, e2)) & (e3 = op(e2, e3)) & (e3 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP59])])).
fof(f114, plain, (((e2 = op(e2, e1)) & (e2 = op(e1, e2)) & (e3 = op(e2, e3)) & (e3 = op(e3, e2))) | ~ sP60), inference(usedef, [], [e114])).
fof(e114, plain, (sP60 <=> ((e2 = op(e2, e1)) & (e2 = op(e1, e2)) & (e3 = op(e2, e3)) & (e3 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP60])])).
fof(f115, plain, (((e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e3 = op(e2, e3)) & (e3 = op(e3, e2))) | ~ sP61), inference(usedef, [], [e115])).
fof(e115, plain, (sP61 <=> ((e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e3 = op(e2, e3)) & (e3 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP61])])).
fof(f116, plain, (((e2 = op(e2, e3)) & (e2 = op(e3, e2)) & (e3 = op(e2, e3)) & (e3 = op(e3, e2))) | ~ sP62), inference(usedef, [], [e116])).
fof(e116, plain, (sP62 <=> ((e2 = op(e2, e3)) & (e2 = op(e3, e2)) & (e3 = op(e2, e3)) & (e3 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP62])])).
fof(f117, plain, (((e3 = op(e3, e0)) & (e3 = op(e0, e3)) & (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ~ sP63), inference(usedef, [], [e117])).
fof(e117, plain, (sP63 <=> ((e3 = op(e3, e0)) & (e3 = op(e0, e3)) & (e3 = op(e3, e3)) & (e3 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP63])])).
fof(f118, plain, (((e3 = op(e3, e1)) & (e3 = op(e1, e3)) & (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ~ sP64), inference(usedef, [], [e118])).
fof(e118, plain, (sP64 <=> ((e3 = op(e3, e1)) & (e3 = op(e1, e3)) & (e3 = op(e3, e3)) & (e3 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP64])])).
fof(f119, plain, (((e3 = op(e3, e2)) & (e3 = op(e2, e3)) & (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ~ sP65), inference(usedef, [], [e119])).
fof(e119, plain, (sP65 <=> ((e3 = op(e3, e2)) & (e3 = op(e2, e3)) & (e3 = op(e3, e3)) & (e3 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP65])])).
fof(f5, plain, ((((e3 = op(e3, e3)) & (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ((e3 = op(e3, e2)) & (e3 = op(e2, e3)) & (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ((e3 = op(e3, e1)) & (e3 = op(e1, e3)) & (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ((e3 = op(e3, e0)) & (e3 = op(e0, e3)) & (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ((e2 = op(e2, e3)) & (e2 = op(e3, e2)) & (e3 = op(e2, e3)) & (e3 = op(e3, e2))) | ((e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e3 = op(e2, e3)) & (e3 = op(e3, e2))) | ((e2 = op(e2, e1)) & (e2 = op(e1, e2)) & (e3 = op(e2, e3)) & (e3 = op(e3, e2))) | ((e2 = op(e2, e0)) & (e2 = op(e0, e2)) & (e3 = op(e2, e3)) & (e3 = op(e3, e2))) | ((e1 = op(e1, e3)) & (e1 = op(e3, e1)) & (e3 = op(e1, e3)) & (e3 = op(e3, e1))) | ((e1 = op(e1, e2)) & (e1 = op(e2, e1)) & (e3 = op(e1, e3)) & (e3 = op(e3, e1))) | ((e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e3 = op(e1, e3)) & (e3 = op(e3, e1))) | ((e1 = op(e1, e0)) & (e1 = op(e0, e1)) & (e3 = op(e1, e3)) & (e3 = op(e3, e1))) | ((e0 = op(e0, e3)) & (e0 = op(e3, e0)) & (e3 = op(e0, e3)) & (e3 = op(e3, e0))) | ((e0 = op(e0, e2)) & (e0 = op(e2, e0)) & (e3 = op(e0, e3)) & (e3 = op(e3, e0))) | ((e0 = op(e0, e1)) & (e0 = op(e1, e0)) & (e3 = op(e0, e3)) & (e3 = op(e3, e0))) | ((e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e3 = op(e0, e3)) & (e3 = op(e3, e0))) | ((e3 = op(e3, e3)) & (e3 = op(e3, e3)) & (e2 = op(e3, e2)) & (e2 = op(e2, e3))) | ((e3 = op(e3, e2)) & (e3 = op(e2, e3)) & (e2 = op(e3, e2)) & (e2 = op(e2, e3))) | ((e3 = op(e3, e1)) & (e3 = op(e1, e3)) & (e2 = op(e3, e2)) & (e2 = op(e2, e3))) | ((e3 = op(e3, e0)) & (e3 = op(e0, e3)) & (e2 = op(e3, e2)) & (e2 = op(e2, e3))) | ((e2 = op(e2, e3)) & (e2 = op(e3, e2)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ((e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ((e2 = op(e2, e1)) & (e2 = op(e1, e2)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ((e2 = op(e2, e0)) & (e2 = op(e0, e2)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ((e1 = op(e1, e3)) & (e1 = op(e3, e1)) & (e2 = op(e1, e2)) & (e2 = op(e2, e1))) | ((e1 = op(e1, e2)) & (e1 = op(e2, e1)) & (e2 = op(e1, e2)) & (e2 = op(e2, e1))) | ((e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e2 = op(e1, e2)) & (e2 = op(e2, e1))) | ((e1 = op(e1, e0)) & (e1 = op(e0, e1)) & (e2 = op(e1, e2)) & (e2 = op(e2, e1))) | ((e0 = op(e0, e3)) & (e0 = op(e3, e0)) & (e2 = op(e0, e2)) & (e2 = op(e2, e0))) | ((e0 = op(e0, e2)) & (e0 = op(e2, e0)) & (e2 = op(e0, e2)) & (e2 = op(e2, e0))) | ((e0 = op(e0, e1)) & (e0 = op(e1, e0)) & (e2 = op(e0, e2)) & (e2 = op(e2, e0))) | ((e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e2 = op(e0, e2)) & (e2 = op(e2, e0))) | ((e3 = op(e3, e3)) & (e3 = op(e3, e3)) & (e1 = op(e3, e1)) & (e1 = op(e1, e3))) | ((e3 = op(e3, e2)) & (e3 = op(e2, e3)) & (e1 = op(e3, e1)) & (e1 = op(e1, e3))) | ((e3 = op(e3, e1)) & (e3 = op(e1, e3)) & (e1 = op(e3, e1)) & (e1 = op(e1, e3))) | ((e3 = op(e3, e0)) & (e3 = op(e0, e3)) & (e1 = op(e3, e1)) & (e1 = op(e1, e3))) | ((e2 = op(e2, e3)) & (e2 = op(e3, e2)) & (e1 = op(e2, e1)) & (e1 = op(e1, e2))) | ((e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e1 = op(e2, e1)) & (e1 = op(e1, e2))) | ((e2 = op(e2, e1)) & (e2 = op(e1, e2)) & (e1 = op(e2, e1)) & (e1 = op(e1, e2))) | ((e2 = op(e2, e0)) & (e2 = op(e0, e2)) & (e1 = op(e2, e1)) & (e1 = op(e1, e2))) | ((e1 = op(e1, e3)) & (e1 = op(e3, e1)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ((e1 = op(e1, e2)) & (e1 = op(e2, e1)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ((e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ((e1 = op(e1, e0)) & (e1 = op(e0, e1)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ((e0 = op(e0, e3)) & (e0 = op(e3, e0)) & (e1 = op(e0, e1)) & (e1 = op(e1, e0))) | ((e0 = op(e0, e2)) & (e0 = op(e2, e0)) & (e1 = op(e0, e1)) & (e1 = op(e1, e0))) | ((e0 = op(e0, e1)) & (e0 = op(e1, e0)) & (e1 = op(e0, e1)) & (e1 = op(e1, e0))) | ((e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e1 = op(e0, e1)) & (e1 = op(e1, e0))) | ((e3 = op(e3, e3)) & (e3 = op(e3, e3)) & (e0 = op(e3, e0)) & (e0 = op(e0, e3))) | ((e3 = op(e3, e2)) & (e3 = op(e2, e3)) & (e0 = op(e3, e0)) & (e0 = op(e0, e3))) | ((e3 = op(e3, e1)) & (e3 = op(e1, e3)) & (e0 = op(e3, e0)) & (e0 = op(e0, e3))) | ((e3 = op(e3, e0)) & (e3 = op(e0, e3)) & (e0 = op(e3, e0)) & (e0 = op(e0, e3))) | ((e2 = op(e2, e3)) & (e2 = op(e3, e2)) & (e0 = op(e2, e0)) & (e0 = op(e0, e2))) | ((e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e0 = op(e2, e0)) & (e0 = op(e0, e2))) | ((e2 = op(e2, e1)) & (e2 = op(e1, e2)) & (e0 = op(e2, e0)) & (e0 = op(e0, e2))) | ((e2 = op(e2, e0)) & (e2 = op(e0, e2)) & (e0 = op(e2, e0)) & (e0 = op(e0, e2))) | ((e1 = op(e1, e3)) & (e1 = op(e3, e1)) & (e0 = op(e1, e0)) & (e0 = op(e0, e1))) | ((e1 = op(e1, e2)) & (e1 = op(e2, e1)) & (e0 = op(e1, e0)) & (e0 = op(e0, e1))) | ((e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e0 = op(e1, e0)) & (e0 = op(e0, e1))) | ((e1 = op(e1, e0)) & (e1 = op(e0, e1)) & (e0 = op(e1, e0)) & (e0 = op(e0, e1))) | ((e0 = op(e0, e3)) & (e0 = op(e3, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ((e0 = op(e0, e2)) & (e0 = op(e2, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ((e0 = op(e0, e1)) & (e0 = op(e1, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ((e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0)))) & ((~ (e3 = op(e3, op(e3, e3))) & ~ (e2 = op(e3, op(e3, e2))) & ~ (e1 = op(e3, op(e3, e1))) & ~ (e0 = op(e3, op(e3, e0)))) | (~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0)))) | (~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0)))) | (~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))))) & ((e3 = op(e3, e3)) | (e2 = op(e2, e3)) | (e1 = op(e1, e3)) | (e0 = op(e0, e3))) & ((e3 = op(e3, e2)) | (e2 = op(e2, e2)) | (e1 = op(e1, e2)) | (e0 = op(e0, e2))) & ((e3 = op(e3, e1)) | (e2 = op(e2, e1)) | (e1 = op(e1, e1)) | (e0 = op(e0, e1))) & ((e3 = op(e3, e0)) | (e2 = op(e2, e0)) | (e1 = op(e1, e0)) | (e0 = op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG143+1.p', ax5)).
fof(f1495, plain, (spl66_57 | spl66_42 | spl66_27 | spl66_12), inference(avatar_split_clause, [], [f554, f636, f700, f764, f828])).
fof(f554, plain, ((e3 = op(e3, e1)) | (e2 = op(e2, e1)) | (e1 = op(e1, e1)) | (e0 = op(e0, e1))), inference(cnf_transformation, [], [f120])).
fof(f1494, plain, (spl66_53 | spl66_38 | spl66_23 | spl66_8), inference(avatar_split_clause, [], [f555, f619, f683, f747, f811])).
fof(f555, plain, ((e3 = op(e3, e2)) | (e2 = op(e2, e2)) | (e1 = op(e1, e2)) | (e0 = op(e0, e2))), inference(cnf_transformation, [], [f120])).
fof(f1493, plain, (spl66_49 | spl66_34 | spl66_19 | spl66_4), inference(avatar_split_clause, [], [f556, f602, f666, f730, f794])).
fof(f556, plain, ((e3 = op(e3, e3)) | (e2 = op(e2, e3)) | (e1 = op(e1, e3)) | (e0 = op(e0, e3))), inference(cnf_transformation, [], [f120])).
fof(f1492, plain, (spl66_138 | spl66_133 | spl66_128 | ~ spl66_146), inference(avatar_split_clause, [], [f557, f1489, f1398, f1422, f1446])).
fof(f1446, plain, (spl66_138 <=> sP0), introduced(avatar_definition, [new_symbols(naming, [spl66_138])])).
fof(f1422, plain, (spl66_133 <=> sP1), introduced(avatar_definition, [new_symbols(naming, [spl66_133])])).
fof(f1398, plain, (spl66_128 <=> sP2), introduced(avatar_definition, [new_symbols(naming, [spl66_128])])).
fof(f557, plain, (~ (e0 = op(e3, op(e3, e0))) | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f120])).
fof(f1477, plain, (spl66_138 | spl66_133 | spl66_128 | ~ spl66_143), inference(avatar_split_clause, [], [f560, f1474, f1398, f1422, f1446])).
fof(f560, plain, (~ (e3 = op(e3, op(e3, e3))) | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f120])).
fof(f1472, plain, (spl66_127 | spl66_126 | spl66_125 | spl66_124 | spl66_123 | spl66_122 | spl66_121 | spl66_120 | spl66_119 | spl66_118 | spl66_117 | spl66_116 | spl66_115 | spl66_114 | spl66_113 | spl66_112 | spl66_111 | spl66_110 | spl66_109 | spl66_108 | spl66_107 | spl66_106 | spl66_105 | spl66_104 | spl66_103 | spl66_102 | spl66_101 | spl66_100 | spl66_99 | spl66_98 | spl66_97 | spl66_96 | spl66_95 | spl66_94 | spl66_93 | spl66_92 | spl66_91 | spl66_90 | spl66_89 | spl66_88 | spl66_87 | spl66_86 | spl66_85 | spl66_84 | spl66_83 | spl66_82 | spl66_81 | spl66_80 | spl66_79 | spl66_78 | spl66_77 | spl66_76 | spl66_75 | spl66_74 | spl66_73 | spl66_72 | spl66_71 | spl66_70 | spl66_69 | spl66_68 | spl66_67 | spl66_66 | spl66_65 | spl66_4), inference(avatar_split_clause, [], [f561, f602, f894, f902, f910, f918, f926, f934, f942, f950, f958, f966, f974, f982, f990, f998, f1006, f1014, f1022, f1030, f1038, f1046, f1054, f1062, f1070, f1078, f1086, f1094, f1102, f1110, f1118, f1126, f1134, f1142, f1150, f1158, f1166, f1174, f1182, f1190, f1198, f1206, f1214, f1222, f1230, f1238, f1246, f1254, f1262, f1270, f1278, f1286, f1294, f1302, f1310, f1318, f1326, f1334, f1342, f1350, f1358, f1366, f1374, f1382, f1390])).
fof(f1390, plain, (spl66_127 <=> sP3), introduced(avatar_definition, [new_symbols(naming, [spl66_127])])).
fof(f1382, plain, (spl66_126 <=> sP4), introduced(avatar_definition, [new_symbols(naming, [spl66_126])])).
fof(f1374, plain, (spl66_125 <=> sP5), introduced(avatar_definition, [new_symbols(naming, [spl66_125])])).
fof(f1366, plain, (spl66_124 <=> sP6), introduced(avatar_definition, [new_symbols(naming, [spl66_124])])).
fof(f1358, plain, (spl66_123 <=> sP7), introduced(avatar_definition, [new_symbols(naming, [spl66_123])])).
fof(f1350, plain, (spl66_122 <=> sP8), introduced(avatar_definition, [new_symbols(naming, [spl66_122])])).
fof(f1342, plain, (spl66_121 <=> sP9), introduced(avatar_definition, [new_symbols(naming, [spl66_121])])).
fof(f1334, plain, (spl66_120 <=> sP10), introduced(avatar_definition, [new_symbols(naming, [spl66_120])])).
fof(f1326, plain, (spl66_119 <=> sP11), introduced(avatar_definition, [new_symbols(naming, [spl66_119])])).
fof(f1318, plain, (spl66_118 <=> sP12), introduced(avatar_definition, [new_symbols(naming, [spl66_118])])).
fof(f1310, plain, (spl66_117 <=> sP13), introduced(avatar_definition, [new_symbols(naming, [spl66_117])])).
fof(f1302, plain, (spl66_116 <=> sP14), introduced(avatar_definition, [new_symbols(naming, [spl66_116])])).
fof(f1294, plain, (spl66_115 <=> sP15), introduced(avatar_definition, [new_symbols(naming, [spl66_115])])).
fof(f1286, plain, (spl66_114 <=> sP16), introduced(avatar_definition, [new_symbols(naming, [spl66_114])])).
fof(f1278, plain, (spl66_113 <=> sP17), introduced(avatar_definition, [new_symbols(naming, [spl66_113])])).
fof(f1270, plain, (spl66_112 <=> sP18), introduced(avatar_definition, [new_symbols(naming, [spl66_112])])).
fof(f1262, plain, (spl66_111 <=> sP19), introduced(avatar_definition, [new_symbols(naming, [spl66_111])])).
fof(f1254, plain, (spl66_110 <=> sP20), introduced(avatar_definition, [new_symbols(naming, [spl66_110])])).
fof(f1246, plain, (spl66_109 <=> sP21), introduced(avatar_definition, [new_symbols(naming, [spl66_109])])).
fof(f1238, plain, (spl66_108 <=> sP22), introduced(avatar_definition, [new_symbols(naming, [spl66_108])])).
fof(f1230, plain, (spl66_107 <=> sP23), introduced(avatar_definition, [new_symbols(naming, [spl66_107])])).
fof(f1222, plain, (spl66_106 <=> sP24), introduced(avatar_definition, [new_symbols(naming, [spl66_106])])).
fof(f1214, plain, (spl66_105 <=> sP25), introduced(avatar_definition, [new_symbols(naming, [spl66_105])])).
fof(f1206, plain, (spl66_104 <=> sP26), introduced(avatar_definition, [new_symbols(naming, [spl66_104])])).
fof(f1198, plain, (spl66_103 <=> sP27), introduced(avatar_definition, [new_symbols(naming, [spl66_103])])).
fof(f1190, plain, (spl66_102 <=> sP28), introduced(avatar_definition, [new_symbols(naming, [spl66_102])])).
fof(f1182, plain, (spl66_101 <=> sP29), introduced(avatar_definition, [new_symbols(naming, [spl66_101])])).
fof(f1174, plain, (spl66_100 <=> sP30), introduced(avatar_definition, [new_symbols(naming, [spl66_100])])).
fof(f1166, plain, (spl66_99 <=> sP31), introduced(avatar_definition, [new_symbols(naming, [spl66_99])])).
fof(f1158, plain, (spl66_98 <=> sP32), introduced(avatar_definition, [new_symbols(naming, [spl66_98])])).
fof(f1150, plain, (spl66_97 <=> sP33), introduced(avatar_definition, [new_symbols(naming, [spl66_97])])).
fof(f1142, plain, (spl66_96 <=> sP34), introduced(avatar_definition, [new_symbols(naming, [spl66_96])])).
fof(f1134, plain, (spl66_95 <=> sP35), introduced(avatar_definition, [new_symbols(naming, [spl66_95])])).
fof(f1126, plain, (spl66_94 <=> sP36), introduced(avatar_definition, [new_symbols(naming, [spl66_94])])).
fof(f1118, plain, (spl66_93 <=> sP37), introduced(avatar_definition, [new_symbols(naming, [spl66_93])])).
fof(f1110, plain, (spl66_92 <=> sP38), introduced(avatar_definition, [new_symbols(naming, [spl66_92])])).
fof(f1102, plain, (spl66_91 <=> sP39), introduced(avatar_definition, [new_symbols(naming, [spl66_91])])).
fof(f1094, plain, (spl66_90 <=> sP40), introduced(avatar_definition, [new_symbols(naming, [spl66_90])])).
fof(f1086, plain, (spl66_89 <=> sP41), introduced(avatar_definition, [new_symbols(naming, [spl66_89])])).
fof(f1078, plain, (spl66_88 <=> sP42), introduced(avatar_definition, [new_symbols(naming, [spl66_88])])).
fof(f1070, plain, (spl66_87 <=> sP43), introduced(avatar_definition, [new_symbols(naming, [spl66_87])])).
fof(f1062, plain, (spl66_86 <=> sP44), introduced(avatar_definition, [new_symbols(naming, [spl66_86])])).
fof(f1054, plain, (spl66_85 <=> sP45), introduced(avatar_definition, [new_symbols(naming, [spl66_85])])).
fof(f1046, plain, (spl66_84 <=> sP46), introduced(avatar_definition, [new_symbols(naming, [spl66_84])])).
fof(f1038, plain, (spl66_83 <=> sP47), introduced(avatar_definition, [new_symbols(naming, [spl66_83])])).
fof(f1030, plain, (spl66_82 <=> sP48), introduced(avatar_definition, [new_symbols(naming, [spl66_82])])).
fof(f1022, plain, (spl66_81 <=> sP49), introduced(avatar_definition, [new_symbols(naming, [spl66_81])])).
fof(f1014, plain, (spl66_80 <=> sP50), introduced(avatar_definition, [new_symbols(naming, [spl66_80])])).
fof(f1006, plain, (spl66_79 <=> sP51), introduced(avatar_definition, [new_symbols(naming, [spl66_79])])).
fof(f998, plain, (spl66_78 <=> sP52), introduced(avatar_definition, [new_symbols(naming, [spl66_78])])).
fof(f990, plain, (spl66_77 <=> sP53), introduced(avatar_definition, [new_symbols(naming, [spl66_77])])).
fof(f982, plain, (spl66_76 <=> sP54), introduced(avatar_definition, [new_symbols(naming, [spl66_76])])).
fof(f974, plain, (spl66_75 <=> sP55), introduced(avatar_definition, [new_symbols(naming, [spl66_75])])).
fof(f966, plain, (spl66_74 <=> sP56), introduced(avatar_definition, [new_symbols(naming, [spl66_74])])).
fof(f958, plain, (spl66_73 <=> sP57), introduced(avatar_definition, [new_symbols(naming, [spl66_73])])).
fof(f950, plain, (spl66_72 <=> sP58), introduced(avatar_definition, [new_symbols(naming, [spl66_72])])).
fof(f942, plain, (spl66_71 <=> sP59), introduced(avatar_definition, [new_symbols(naming, [spl66_71])])).
fof(f934, plain, (spl66_70 <=> sP60), introduced(avatar_definition, [new_symbols(naming, [spl66_70])])).
fof(f926, plain, (spl66_69 <=> sP61), introduced(avatar_definition, [new_symbols(naming, [spl66_69])])).
fof(f918, plain, (spl66_68 <=> sP62), introduced(avatar_definition, [new_symbols(naming, [spl66_68])])).
fof(f910, plain, (spl66_67 <=> sP63), introduced(avatar_definition, [new_symbols(naming, [spl66_67])])).
fof(f902, plain, (spl66_66 <=> sP64), introduced(avatar_definition, [new_symbols(naming, [spl66_66])])).
fof(f894, plain, (spl66_65 <=> sP65), introduced(avatar_definition, [new_symbols(naming, [spl66_65])])).
fof(f561, plain, ((e3 = op(e3, e3)) | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3), inference(cnf_transformation, [], [f120])).
fof(f1468, plain, (~ spl66_138 | ~ spl66_142), inference(avatar_split_clause, [], [f549, f1465, f1446])).
fof(f549, plain, (~ (e0 = op(e0, op(e0, e0))) | ~ sP0), inference(cnf_transformation, [], [f186])).
fof(f186, plain, ((~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0)))) | ~ sP0), inference(nnf_transformation, [], [f54])).
fof(f1463, plain, (~ spl66_138 | ~ spl66_141), inference(avatar_split_clause, [], [f550, f1460, f1446])).
fof(f550, plain, (~ (e1 = op(e0, op(e0, e1))) | ~ sP0), inference(cnf_transformation, [], [f186])).
fof(f1458, plain, (~ spl66_138 | ~ spl66_140), inference(avatar_split_clause, [], [f551, f1455, f1446])).
fof(f551, plain, (~ (e2 = op(e0, op(e0, e2))) | ~ sP0), inference(cnf_transformation, [], [f186])).
fof(f1444, plain, (~ spl66_133 | ~ spl66_137), inference(avatar_split_clause, [], [f545, f1441, f1422])).
fof(f545, plain, (~ (e0 = op(e1, op(e1, e0))) | ~ sP1), inference(cnf_transformation, [], [f185])).
fof(f185, plain, ((~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0)))) | ~ sP1), inference(nnf_transformation, [], [f55])).
fof(f1439, plain, (~ spl66_133 | ~ spl66_136), inference(avatar_split_clause, [], [f546, f1436, f1422])).
fof(f546, plain, (~ (e1 = op(e1, op(e1, e1))) | ~ sP1), inference(cnf_transformation, [], [f185])).
fof(f1434, plain, (~ spl66_133 | ~ spl66_135), inference(avatar_split_clause, [], [f547, f1431, f1422])).
fof(f547, plain, (~ (e2 = op(e1, op(e1, e2))) | ~ sP1), inference(cnf_transformation, [], [f185])).
fof(f1420, plain, (~ spl66_128 | ~ spl66_132), inference(avatar_split_clause, [], [f541, f1417, f1398])).
fof(f541, plain, (~ (e0 = op(e2, op(e2, e0))) | ~ sP2), inference(cnf_transformation, [], [f184])).
fof(f184, plain, ((~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0)))) | ~ sP2), inference(nnf_transformation, [], [f56])).
fof(f1415, plain, (~ spl66_128 | ~ spl66_131), inference(avatar_split_clause, [], [f542, f1412, f1398])).
fof(f542, plain, (~ (e1 = op(e2, op(e2, e1))) | ~ sP2), inference(cnf_transformation, [], [f184])).
fof(f1410, plain, (~ spl66_128 | ~ spl66_130), inference(avatar_split_clause, [], [f543, f1407, f1398])).
fof(f543, plain, (~ (e2 = op(e2, op(e2, e2))) | ~ sP2), inference(cnf_transformation, [], [f184])).
fof(f1405, plain, (~ spl66_128 | ~ spl66_129), inference(avatar_split_clause, [], [f544, f1402, f1398])).
fof(f544, plain, (~ (e3 = op(e2, op(e2, e3))) | ~ sP2), inference(cnf_transformation, [], [f184])).
fof(f1396, plain, (~ spl66_127 | spl66_61), inference(avatar_split_clause, [], [f537, f845, f1390])).
fof(f537, plain, ((e0 = op(e0, e0)) | ~ sP3), inference(cnf_transformation, [], [f183])).
fof(f183, plain, (((e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP3), inference(nnf_transformation, [], [f57])).
fof(f1388, plain, (~ spl66_126 | spl66_61), inference(avatar_split_clause, [], [f533, f845, f1382])).
fof(f533, plain, ((e0 = op(e0, e0)) | ~ sP4), inference(cnf_transformation, [], [f182])).
fof(f182, plain, (((e0 = op(e0, e1)) & (e0 = op(e1, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP4), inference(nnf_transformation, [], [f58])).
fof(f1385, plain, (~ spl66_126 | spl66_57), inference(avatar_split_clause, [], [f536, f828, f1382])).
fof(f536, plain, ((e0 = op(e0, e1)) | ~ sP4), inference(cnf_transformation, [], [f182])).
fof(f1380, plain, (~ spl66_125 | spl66_61), inference(avatar_split_clause, [], [f529, f845, f1374])).
fof(f529, plain, ((e0 = op(e0, e0)) | ~ sP5), inference(cnf_transformation, [], [f181])).
fof(f181, plain, (((e0 = op(e0, e2)) & (e0 = op(e2, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP5), inference(nnf_transformation, [], [f59])).
fof(f1372, plain, (~ spl66_124 | spl66_61), inference(avatar_split_clause, [], [f525, f845, f1366])).
fof(f525, plain, ((e0 = op(e0, e0)) | ~ sP6), inference(cnf_transformation, [], [f180])).
fof(f180, plain, (((e0 = op(e0, e3)) & (e0 = op(e3, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP6), inference(nnf_transformation, [], [f60])).
fof(f1364, plain, (~ spl66_123 | spl66_57), inference(avatar_split_clause, [], [f521, f828, f1358])).
fof(f521, plain, ((e0 = op(e0, e1)) | ~ sP7), inference(cnf_transformation, [], [f179])).
fof(f179, plain, (((e1 = op(e1, e0)) & (e1 = op(e0, e1)) & (e0 = op(e1, e0)) & (e0 = op(e0, e1))) | ~ sP7), inference(nnf_transformation, [], [f61])).
fof(f1363, plain, (~ spl66_123 | spl66_45), inference(avatar_split_clause, [], [f522, f777, f1358])).
fof(f522, plain, ((e0 = op(e1, e0)) | ~ sP7), inference(cnf_transformation, [], [f179])).
fof(f1361, plain, (~ spl66_123 | spl66_46), inference(avatar_split_clause, [], [f524, f781, f1358])).
fof(f524, plain, ((e1 = op(e1, e0)) | ~ sP7), inference(cnf_transformation, [], [f179])).
fof(f1356, plain, (~ spl66_122 | spl66_57), inference(avatar_split_clause, [], [f517, f828, f1350])).
fof(f517, plain, ((e0 = op(e0, e1)) | ~ sP8), inference(cnf_transformation, [], [f178])).
fof(f178, plain, (((e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e0 = op(e1, e0)) & (e0 = op(e0, e1))) | ~ sP8), inference(nnf_transformation, [], [f62])).
fof(f1354, plain, (~ spl66_122 | spl66_42), inference(avatar_split_clause, [], [f519, f764, f1350])).
fof(f519, plain, ((e1 = op(e1, e1)) | ~ sP8), inference(cnf_transformation, [], [f178])).
fof(f1348, plain, (~ spl66_121 | spl66_57), inference(avatar_split_clause, [], [f513, f828, f1342])).
fof(f513, plain, ((e0 = op(e0, e1)) | ~ sP9), inference(cnf_transformation, [], [f177])).
fof(f177, plain, (((e1 = op(e1, e2)) & (e1 = op(e2, e1)) & (e0 = op(e1, e0)) & (e0 = op(e0, e1))) | ~ sP9), inference(nnf_transformation, [], [f63])).
fof(f1347, plain, (~ spl66_121 | spl66_45), inference(avatar_split_clause, [], [f514, f777, f1342])).
fof(f514, plain, ((e0 = op(e1, e0)) | ~ sP9), inference(cnf_transformation, [], [f177])).
fof(f1346, plain, (~ spl66_121 | spl66_26), inference(avatar_split_clause, [], [f515, f696, f1342])).
fof(f515, plain, ((e1 = op(e2, e1)) | ~ sP9), inference(cnf_transformation, [], [f177])).
fof(f1345, plain, (~ spl66_121 | spl66_38), inference(avatar_split_clause, [], [f516, f747, f1342])).
fof(f516, plain, ((e1 = op(e1, e2)) | ~ sP9), inference(cnf_transformation, [], [f177])).
fof(f1340, plain, (~ spl66_120 | spl66_57), inference(avatar_split_clause, [], [f509, f828, f1334])).
fof(f509, plain, ((e0 = op(e0, e1)) | ~ sP10), inference(cnf_transformation, [], [f176])).
fof(f176, plain, (((e1 = op(e1, e3)) & (e1 = op(e3, e1)) & (e0 = op(e1, e0)) & (e0 = op(e0, e1))) | ~ sP10), inference(nnf_transformation, [], [f64])).
fof(f1339, plain, (~ spl66_120 | spl66_45), inference(avatar_split_clause, [], [f510, f777, f1334])).
fof(f510, plain, ((e0 = op(e1, e0)) | ~ sP10), inference(cnf_transformation, [], [f176])).
fof(f1338, plain, (~ spl66_120 | spl66_10), inference(avatar_split_clause, [], [f511, f628, f1334])).
fof(f511, plain, ((e1 = op(e3, e1)) | ~ sP10), inference(cnf_transformation, [], [f176])).
fof(f1337, plain, (~ spl66_120 | spl66_34), inference(avatar_split_clause, [], [f512, f730, f1334])).
fof(f512, plain, ((e1 = op(e1, e3)) | ~ sP10), inference(cnf_transformation, [], [f176])).
fof(f1332, plain, (~ spl66_119 | spl66_53), inference(avatar_split_clause, [], [f505, f811, f1326])).
fof(f505, plain, ((e0 = op(e0, e2)) | ~ sP11), inference(cnf_transformation, [], [f175])).
fof(f175, plain, (((e2 = op(e2, e0)) & (e2 = op(e0, e2)) & (e0 = op(e2, e0)) & (e0 = op(e0, e2))) | ~ sP11), inference(nnf_transformation, [], [f65])).
fof(f1331, plain, (~ spl66_119 | spl66_29), inference(avatar_split_clause, [], [f506, f709, f1326])).
fof(f506, plain, ((e0 = op(e2, e0)) | ~ sP11), inference(cnf_transformation, [], [f175])).
fof(f1330, plain, (~ spl66_119 | spl66_55), inference(avatar_split_clause, [], [f507, f819, f1326])).
fof(f507, plain, ((e2 = op(e0, e2)) | ~ sP11), inference(cnf_transformation, [], [f175])).
fof(f1329, plain, (~ spl66_119 | spl66_31), inference(avatar_split_clause, [], [f508, f717, f1326])).
fof(f508, plain, ((e2 = op(e2, e0)) | ~ sP11), inference(cnf_transformation, [], [f175])).
fof(f1324, plain, (~ spl66_118 | spl66_53), inference(avatar_split_clause, [], [f501, f811, f1318])).
fof(f501, plain, ((e0 = op(e0, e2)) | ~ sP12), inference(cnf_transformation, [], [f174])).
fof(f174, plain, (((e2 = op(e2, e1)) & (e2 = op(e1, e2)) & (e0 = op(e2, e0)) & (e0 = op(e0, e2))) | ~ sP12), inference(nnf_transformation, [], [f66])).
fof(f1323, plain, (~ spl66_118 | spl66_29), inference(avatar_split_clause, [], [f502, f709, f1318])).
fof(f502, plain, ((e0 = op(e2, e0)) | ~ sP12), inference(cnf_transformation, [], [f174])).
fof(f1322, plain, (~ spl66_118 | spl66_39), inference(avatar_split_clause, [], [f503, f751, f1318])).
fof(f503, plain, ((e2 = op(e1, e2)) | ~ sP12), inference(cnf_transformation, [], [f174])).
fof(f1321, plain, (~ spl66_118 | spl66_27), inference(avatar_split_clause, [], [f504, f700, f1318])).
fof(f504, plain, ((e2 = op(e2, e1)) | ~ sP12), inference(cnf_transformation, [], [f174])).
fof(f1316, plain, (~ spl66_117 | spl66_53), inference(avatar_split_clause, [], [f497, f811, f1310])).
fof(f497, plain, ((e0 = op(e0, e2)) | ~ sP13), inference(cnf_transformation, [], [f173])).
fof(f173, plain, (((e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e0 = op(e2, e0)) & (e0 = op(e0, e2))) | ~ sP13), inference(nnf_transformation, [], [f67])).
fof(f1314, plain, (~ spl66_117 | spl66_23), inference(avatar_split_clause, [], [f499, f683, f1310])).
fof(f499, plain, ((e2 = op(e2, e2)) | ~ sP13), inference(cnf_transformation, [], [f173])).
fof(f1308, plain, (~ spl66_116 | spl66_53), inference(avatar_split_clause, [], [f493, f811, f1302])).
fof(f493, plain, ((e0 = op(e0, e2)) | ~ sP14), inference(cnf_transformation, [], [f172])).
fof(f172, plain, (((e2 = op(e2, e3)) & (e2 = op(e3, e2)) & (e0 = op(e2, e0)) & (e0 = op(e0, e2))) | ~ sP14), inference(nnf_transformation, [], [f68])).
fof(f1307, plain, (~ spl66_116 | spl66_29), inference(avatar_split_clause, [], [f494, f709, f1302])).
fof(f494, plain, ((e0 = op(e2, e0)) | ~ sP14), inference(cnf_transformation, [], [f172])).
fof(f1306, plain, (~ spl66_116 | spl66_7), inference(avatar_split_clause, [], [f495, f615, f1302])).
fof(f495, plain, ((e2 = op(e3, e2)) | ~ sP14), inference(cnf_transformation, [], [f172])).
fof(f1305, plain, (~ spl66_116 | spl66_19), inference(avatar_split_clause, [], [f496, f666, f1302])).
fof(f496, plain, ((e2 = op(e2, e3)) | ~ sP14), inference(cnf_transformation, [], [f172])).
fof(f1300, plain, (~ spl66_115 | spl66_49), inference(avatar_split_clause, [], [f489, f794, f1294])).
fof(f489, plain, ((e0 = op(e0, e3)) | ~ sP15), inference(cnf_transformation, [], [f171])).
fof(f171, plain, (((e3 = op(e3, e0)) & (e3 = op(e0, e3)) & (e0 = op(e3, e0)) & (e0 = op(e0, e3))) | ~ sP15), inference(nnf_transformation, [], [f69])).
fof(f1299, plain, (~ spl66_115 | spl66_13), inference(avatar_split_clause, [], [f490, f641, f1294])).
fof(f490, plain, ((e0 = op(e3, e0)) | ~ sP15), inference(cnf_transformation, [], [f171])).
fof(f1298, plain, (~ spl66_115 | spl66_52), inference(avatar_split_clause, [], [f491, f806, f1294])).
fof(f491, plain, ((e3 = op(e0, e3)) | ~ sP15), inference(cnf_transformation, [], [f171])).
fof(f1297, plain, (~ spl66_115 | spl66_16), inference(avatar_split_clause, [], [f492, f653, f1294])).
fof(f492, plain, ((e3 = op(e3, e0)) | ~ sP15), inference(cnf_transformation, [], [f171])).
fof(f1292, plain, (~ spl66_114 | spl66_49), inference(avatar_split_clause, [], [f485, f794, f1286])).
fof(f485, plain, ((e0 = op(e0, e3)) | ~ sP16), inference(cnf_transformation, [], [f170])).
fof(f170, plain, (((e3 = op(e3, e1)) & (e3 = op(e1, e3)) & (e0 = op(e3, e0)) & (e0 = op(e0, e3))) | ~ sP16), inference(nnf_transformation, [], [f70])).
fof(f1291, plain, (~ spl66_114 | spl66_13), inference(avatar_split_clause, [], [f486, f641, f1286])).
fof(f486, plain, ((e0 = op(e3, e0)) | ~ sP16), inference(cnf_transformation, [], [f170])).
fof(f1290, plain, (~ spl66_114 | spl66_36), inference(avatar_split_clause, [], [f487, f738, f1286])).
fof(f487, plain, ((e3 = op(e1, e3)) | ~ sP16), inference(cnf_transformation, [], [f170])).
fof(f1289, plain, (~ spl66_114 | spl66_12), inference(avatar_split_clause, [], [f488, f636, f1286])).
fof(f488, plain, ((e3 = op(e3, e1)) | ~ sP16), inference(cnf_transformation, [], [f170])).
fof(f1284, plain, (~ spl66_113 | spl66_49), inference(avatar_split_clause, [], [f481, f794, f1278])).
fof(f481, plain, ((e0 = op(e0, e3)) | ~ sP17), inference(cnf_transformation, [], [f169])).
fof(f169, plain, (((e3 = op(e3, e2)) & (e3 = op(e2, e3)) & (e0 = op(e3, e0)) & (e0 = op(e0, e3))) | ~ sP17), inference(nnf_transformation, [], [f71])).
fof(f1283, plain, (~ spl66_113 | spl66_13), inference(avatar_split_clause, [], [f482, f641, f1278])).
fof(f482, plain, ((e0 = op(e3, e0)) | ~ sP17), inference(cnf_transformation, [], [f169])).
fof(f1282, plain, (~ spl66_113 | spl66_20), inference(avatar_split_clause, [], [f483, f670, f1278])).
fof(f483, plain, ((e3 = op(e2, e3)) | ~ sP17), inference(cnf_transformation, [], [f169])).
fof(f1281, plain, (~ spl66_113 | spl66_8), inference(avatar_split_clause, [], [f484, f619, f1278])).
fof(f484, plain, ((e3 = op(e3, e2)) | ~ sP17), inference(cnf_transformation, [], [f169])).
fof(f1276, plain, (~ spl66_112 | spl66_49), inference(avatar_split_clause, [], [f477, f794, f1270])).
fof(f477, plain, ((e0 = op(e0, e3)) | ~ sP18), inference(cnf_transformation, [], [f168])).
fof(f168, plain, (((e3 = op(e3, e3)) & (e3 = op(e3, e3)) & (e0 = op(e3, e0)) & (e0 = op(e0, e3))) | ~ sP18), inference(nnf_transformation, [], [f72])).
fof(f1274, plain, (~ spl66_112 | spl66_4), inference(avatar_split_clause, [], [f479, f602, f1270])).
fof(f479, plain, ((e3 = op(e3, e3)) | ~ sP18), inference(cnf_transformation, [], [f168])).
fof(f1266, plain, (~ spl66_111 | spl66_61), inference(avatar_split_clause, [], [f475, f845, f1262])).
fof(f475, plain, ((e0 = op(e0, e0)) | ~ sP19), inference(cnf_transformation, [], [f167])).
fof(f167, plain, (((e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e1 = op(e0, e1)) & (e1 = op(e1, e0))) | ~ sP19), inference(nnf_transformation, [], [f73])).
fof(f1260, plain, (~ spl66_110 | spl66_46), inference(avatar_split_clause, [], [f469, f781, f1254])).
fof(f469, plain, ((e1 = op(e1, e0)) | ~ sP20), inference(cnf_transformation, [], [f166])).
fof(f166, plain, (((e0 = op(e0, e1)) & (e0 = op(e1, e0)) & (e1 = op(e0, e1)) & (e1 = op(e1, e0))) | ~ sP20), inference(nnf_transformation, [], [f74])).
fof(f1258, plain, (~ spl66_110 | spl66_45), inference(avatar_split_clause, [], [f471, f777, f1254])).
fof(f471, plain, ((e0 = op(e1, e0)) | ~ sP20), inference(cnf_transformation, [], [f166])).
fof(f1257, plain, (~ spl66_110 | spl66_57), inference(avatar_split_clause, [], [f472, f828, f1254])).
fof(f472, plain, ((e0 = op(e0, e1)) | ~ sP20), inference(cnf_transformation, [], [f166])).
fof(f1252, plain, (~ spl66_109 | spl66_46), inference(avatar_split_clause, [], [f465, f781, f1246])).
fof(f465, plain, ((e1 = op(e1, e0)) | ~ sP21), inference(cnf_transformation, [], [f165])).
fof(f165, plain, (((e0 = op(e0, e2)) & (e0 = op(e2, e0)) & (e1 = op(e0, e1)) & (e1 = op(e1, e0))) | ~ sP21), inference(nnf_transformation, [], [f75])).
fof(f1251, plain, (~ spl66_109 | spl66_58), inference(avatar_split_clause, [], [f466, f832, f1246])).
fof(f466, plain, ((e1 = op(e0, e1)) | ~ sP21), inference(cnf_transformation, [], [f165])).
fof(f1250, plain, (~ spl66_109 | spl66_29), inference(avatar_split_clause, [], [f467, f709, f1246])).
fof(f467, plain, ((e0 = op(e2, e0)) | ~ sP21), inference(cnf_transformation, [], [f165])).
fof(f1249, plain, (~ spl66_109 | spl66_53), inference(avatar_split_clause, [], [f468, f811, f1246])).
fof(f468, plain, ((e0 = op(e0, e2)) | ~ sP21), inference(cnf_transformation, [], [f165])).
fof(f1244, plain, (~ spl66_108 | spl66_46), inference(avatar_split_clause, [], [f461, f781, f1238])).
fof(f461, plain, ((e1 = op(e1, e0)) | ~ sP22), inference(cnf_transformation, [], [f164])).
fof(f164, plain, (((e0 = op(e0, e3)) & (e0 = op(e3, e0)) & (e1 = op(e0, e1)) & (e1 = op(e1, e0))) | ~ sP22), inference(nnf_transformation, [], [f76])).
fof(f1243, plain, (~ spl66_108 | spl66_58), inference(avatar_split_clause, [], [f462, f832, f1238])).
fof(f462, plain, ((e1 = op(e0, e1)) | ~ sP22), inference(cnf_transformation, [], [f164])).
fof(f1242, plain, (~ spl66_108 | spl66_13), inference(avatar_split_clause, [], [f463, f641, f1238])).
fof(f463, plain, ((e0 = op(e3, e0)) | ~ sP22), inference(cnf_transformation, [], [f164])).
fof(f1241, plain, (~ spl66_108 | spl66_49), inference(avatar_split_clause, [], [f464, f794, f1238])).
fof(f464, plain, ((e0 = op(e0, e3)) | ~ sP22), inference(cnf_transformation, [], [f164])).
fof(f1236, plain, (~ spl66_107 | spl66_42), inference(avatar_split_clause, [], [f457, f764, f1230])).
fof(f457, plain, ((e1 = op(e1, e1)) | ~ sP23), inference(cnf_transformation, [], [f163])).
fof(f163, plain, (((e1 = op(e1, e0)) & (e1 = op(e0, e1)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP23), inference(nnf_transformation, [], [f77])).
fof(f1233, plain, (~ spl66_107 | spl66_46), inference(avatar_split_clause, [], [f460, f781, f1230])).
fof(f460, plain, ((e1 = op(e1, e0)) | ~ sP23), inference(cnf_transformation, [], [f163])).
fof(f1228, plain, (~ spl66_106 | spl66_42), inference(avatar_split_clause, [], [f453, f764, f1222])).
fof(f453, plain, ((e1 = op(e1, e1)) | ~ sP24), inference(cnf_transformation, [], [f162])).
fof(f162, plain, (((e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP24), inference(nnf_transformation, [], [f78])).
fof(f1220, plain, (~ spl66_105 | spl66_42), inference(avatar_split_clause, [], [f449, f764, f1214])).
fof(f449, plain, ((e1 = op(e1, e1)) | ~ sP25), inference(cnf_transformation, [], [f161])).
fof(f161, plain, (((e1 = op(e1, e2)) & (e1 = op(e2, e1)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP25), inference(nnf_transformation, [], [f79])).
fof(f1217, plain, (~ spl66_105 | spl66_38), inference(avatar_split_clause, [], [f452, f747, f1214])).
fof(f452, plain, ((e1 = op(e1, e2)) | ~ sP25), inference(cnf_transformation, [], [f161])).
fof(f1212, plain, (~ spl66_104 | spl66_42), inference(avatar_split_clause, [], [f445, f764, f1206])).
fof(f445, plain, ((e1 = op(e1, e1)) | ~ sP26), inference(cnf_transformation, [], [f160])).
fof(f160, plain, (((e1 = op(e1, e3)) & (e1 = op(e3, e1)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP26), inference(nnf_transformation, [], [f80])).
fof(f1210, plain, (~ spl66_104 | spl66_10), inference(avatar_split_clause, [], [f447, f628, f1206])).
fof(f447, plain, ((e1 = op(e3, e1)) | ~ sP26), inference(cnf_transformation, [], [f160])).
fof(f1209, plain, (~ spl66_104 | spl66_34), inference(avatar_split_clause, [], [f448, f730, f1206])).
fof(f448, plain, ((e1 = op(e1, e3)) | ~ sP26), inference(cnf_transformation, [], [f160])).
fof(f1204, plain, (~ spl66_103 | spl66_38), inference(avatar_split_clause, [], [f441, f747, f1198])).
fof(f441, plain, ((e1 = op(e1, e2)) | ~ sP27), inference(cnf_transformation, [], [f159])).
fof(f159, plain, (((e2 = op(e2, e0)) & (e2 = op(e0, e2)) & (e1 = op(e2, e1)) & (e1 = op(e1, e2))) | ~ sP27), inference(nnf_transformation, [], [f81])).
fof(f1203, plain, (~ spl66_103 | spl66_26), inference(avatar_split_clause, [], [f442, f696, f1198])).
fof(f442, plain, ((e1 = op(e2, e1)) | ~ sP27), inference(cnf_transformation, [], [f159])).
fof(f1202, plain, (~ spl66_103 | spl66_55), inference(avatar_split_clause, [], [f443, f819, f1198])).
fof(f443, plain, ((e2 = op(e0, e2)) | ~ sP27), inference(cnf_transformation, [], [f159])).
fof(f1201, plain, (~ spl66_103 | spl66_31), inference(avatar_split_clause, [], [f444, f717, f1198])).
fof(f444, plain, ((e2 = op(e2, e0)) | ~ sP27), inference(cnf_transformation, [], [f159])).
fof(f1196, plain, (~ spl66_102 | spl66_38), inference(avatar_split_clause, [], [f437, f747, f1190])).
fof(f437, plain, ((e1 = op(e1, e2)) | ~ sP28), inference(cnf_transformation, [], [f158])).
fof(f158, plain, (((e2 = op(e2, e1)) & (e2 = op(e1, e2)) & (e1 = op(e2, e1)) & (e1 = op(e1, e2))) | ~ sP28), inference(nnf_transformation, [], [f82])).
fof(f1195, plain, (~ spl66_102 | spl66_26), inference(avatar_split_clause, [], [f438, f696, f1190])).
fof(f438, plain, ((e1 = op(e2, e1)) | ~ sP28), inference(cnf_transformation, [], [f158])).
fof(f1194, plain, (~ spl66_102 | spl66_39), inference(avatar_split_clause, [], [f439, f751, f1190])).
fof(f439, plain, ((e2 = op(e1, e2)) | ~ sP28), inference(cnf_transformation, [], [f158])).
fof(f1193, plain, (~ spl66_102 | spl66_27), inference(avatar_split_clause, [], [f440, f700, f1190])).
fof(f440, plain, ((e2 = op(e2, e1)) | ~ sP28), inference(cnf_transformation, [], [f158])).
fof(f1188, plain, (~ spl66_101 | spl66_38), inference(avatar_split_clause, [], [f433, f747, f1182])).
fof(f433, plain, ((e1 = op(e1, e2)) | ~ sP29), inference(cnf_transformation, [], [f157])).
fof(f157, plain, (((e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e1 = op(e2, e1)) & (e1 = op(e1, e2))) | ~ sP29), inference(nnf_transformation, [], [f83])).
fof(f1186, plain, (~ spl66_101 | spl66_23), inference(avatar_split_clause, [], [f435, f683, f1182])).
fof(f435, plain, ((e2 = op(e2, e2)) | ~ sP29), inference(cnf_transformation, [], [f157])).
fof(f1180, plain, (~ spl66_100 | spl66_38), inference(avatar_split_clause, [], [f429, f747, f1174])).
fof(f429, plain, ((e1 = op(e1, e2)) | ~ sP30), inference(cnf_transformation, [], [f156])).
fof(f156, plain, (((e2 = op(e2, e3)) & (e2 = op(e3, e2)) & (e1 = op(e2, e1)) & (e1 = op(e1, e2))) | ~ sP30), inference(nnf_transformation, [], [f84])).
fof(f1179, plain, (~ spl66_100 | spl66_26), inference(avatar_split_clause, [], [f430, f696, f1174])).
fof(f430, plain, ((e1 = op(e2, e1)) | ~ sP30), inference(cnf_transformation, [], [f156])).
fof(f1178, plain, (~ spl66_100 | spl66_7), inference(avatar_split_clause, [], [f431, f615, f1174])).
fof(f431, plain, ((e2 = op(e3, e2)) | ~ sP30), inference(cnf_transformation, [], [f156])).
fof(f1177, plain, (~ spl66_100 | spl66_19), inference(avatar_split_clause, [], [f432, f666, f1174])).
fof(f432, plain, ((e2 = op(e2, e3)) | ~ sP30), inference(cnf_transformation, [], [f156])).
fof(f1172, plain, (~ spl66_99 | spl66_34), inference(avatar_split_clause, [], [f425, f730, f1166])).
fof(f425, plain, ((e1 = op(e1, e3)) | ~ sP31), inference(cnf_transformation, [], [f155])).
fof(f155, plain, (((e3 = op(e3, e0)) & (e3 = op(e0, e3)) & (e1 = op(e3, e1)) & (e1 = op(e1, e3))) | ~ sP31), inference(nnf_transformation, [], [f85])).
fof(f1171, plain, (~ spl66_99 | spl66_10), inference(avatar_split_clause, [], [f426, f628, f1166])).
fof(f426, plain, ((e1 = op(e3, e1)) | ~ sP31), inference(cnf_transformation, [], [f155])).
fof(f1170, plain, (~ spl66_99 | spl66_52), inference(avatar_split_clause, [], [f427, f806, f1166])).
fof(f427, plain, ((e3 = op(e0, e3)) | ~ sP31), inference(cnf_transformation, [], [f155])).
fof(f1169, plain, (~ spl66_99 | spl66_16), inference(avatar_split_clause, [], [f428, f653, f1166])).
fof(f428, plain, ((e3 = op(e3, e0)) | ~ sP31), inference(cnf_transformation, [], [f155])).
fof(f1164, plain, (~ spl66_98 | spl66_34), inference(avatar_split_clause, [], [f421, f730, f1158])).
fof(f421, plain, ((e1 = op(e1, e3)) | ~ sP32), inference(cnf_transformation, [], [f154])).
fof(f154, plain, (((e3 = op(e3, e1)) & (e3 = op(e1, e3)) & (e1 = op(e3, e1)) & (e1 = op(e1, e3))) | ~ sP32), inference(nnf_transformation, [], [f86])).
fof(f1163, plain, (~ spl66_98 | spl66_10), inference(avatar_split_clause, [], [f422, f628, f1158])).
fof(f422, plain, ((e1 = op(e3, e1)) | ~ sP32), inference(cnf_transformation, [], [f154])).
fof(f1162, plain, (~ spl66_98 | spl66_36), inference(avatar_split_clause, [], [f423, f738, f1158])).
fof(f423, plain, ((e3 = op(e1, e3)) | ~ sP32), inference(cnf_transformation, [], [f154])).
fof(f1161, plain, (~ spl66_98 | spl66_12), inference(avatar_split_clause, [], [f424, f636, f1158])).
fof(f424, plain, ((e3 = op(e3, e1)) | ~ sP32), inference(cnf_transformation, [], [f154])).
fof(f1156, plain, (~ spl66_97 | spl66_34), inference(avatar_split_clause, [], [f417, f730, f1150])).
fof(f417, plain, ((e1 = op(e1, e3)) | ~ sP33), inference(cnf_transformation, [], [f153])).
fof(f153, plain, (((e3 = op(e3, e2)) & (e3 = op(e2, e3)) & (e1 = op(e3, e1)) & (e1 = op(e1, e3))) | ~ sP33), inference(nnf_transformation, [], [f87])).
fof(f1155, plain, (~ spl66_97 | spl66_10), inference(avatar_split_clause, [], [f418, f628, f1150])).
fof(f418, plain, ((e1 = op(e3, e1)) | ~ sP33), inference(cnf_transformation, [], [f153])).
fof(f1154, plain, (~ spl66_97 | spl66_20), inference(avatar_split_clause, [], [f419, f670, f1150])).
fof(f419, plain, ((e3 = op(e2, e3)) | ~ sP33), inference(cnf_transformation, [], [f153])).
fof(f1153, plain, (~ spl66_97 | spl66_8), inference(avatar_split_clause, [], [f420, f619, f1150])).
fof(f420, plain, ((e3 = op(e3, e2)) | ~ sP33), inference(cnf_transformation, [], [f153])).
fof(f1148, plain, (~ spl66_96 | spl66_34), inference(avatar_split_clause, [], [f413, f730, f1142])).
fof(f413, plain, ((e1 = op(e1, e3)) | ~ sP34), inference(cnf_transformation, [], [f152])).
fof(f152, plain, (((e3 = op(e3, e3)) & (e3 = op(e3, e3)) & (e1 = op(e3, e1)) & (e1 = op(e1, e3))) | ~ sP34), inference(nnf_transformation, [], [f88])).
fof(f1146, plain, (~ spl66_96 | spl66_4), inference(avatar_split_clause, [], [f415, f602, f1142])).
fof(f415, plain, ((e3 = op(e3, e3)) | ~ sP34), inference(cnf_transformation, [], [f152])).
fof(f1138, plain, (~ spl66_95 | spl66_61), inference(avatar_split_clause, [], [f411, f845, f1134])).
fof(f411, plain, ((e0 = op(e0, e0)) | ~ sP35), inference(cnf_transformation, [], [f151])).
fof(f151, plain, (((e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e2 = op(e0, e2)) & (e2 = op(e2, e0))) | ~ sP35), inference(nnf_transformation, [], [f89])).
fof(f1132, plain, (~ spl66_94 | spl66_31), inference(avatar_split_clause, [], [f405, f717, f1126])).
fof(f405, plain, ((e2 = op(e2, e0)) | ~ sP36), inference(cnf_transformation, [], [f150])).
fof(f150, plain, (((e0 = op(e0, e1)) & (e0 = op(e1, e0)) & (e2 = op(e0, e2)) & (e2 = op(e2, e0))) | ~ sP36), inference(nnf_transformation, [], [f90])).
fof(f1131, plain, (~ spl66_94 | spl66_55), inference(avatar_split_clause, [], [f406, f819, f1126])).
fof(f406, plain, ((e2 = op(e0, e2)) | ~ sP36), inference(cnf_transformation, [], [f150])).
fof(f1130, plain, (~ spl66_94 | spl66_45), inference(avatar_split_clause, [], [f407, f777, f1126])).
fof(f407, plain, ((e0 = op(e1, e0)) | ~ sP36), inference(cnf_transformation, [], [f150])).
fof(f1129, plain, (~ spl66_94 | spl66_57), inference(avatar_split_clause, [], [f408, f828, f1126])).
fof(f408, plain, ((e0 = op(e0, e1)) | ~ sP36), inference(cnf_transformation, [], [f150])).
fof(f1124, plain, (~ spl66_93 | spl66_31), inference(avatar_split_clause, [], [f401, f717, f1118])).
fof(f401, plain, ((e2 = op(e2, e0)) | ~ sP37), inference(cnf_transformation, [], [f149])).
fof(f149, plain, (((e0 = op(e0, e2)) & (e0 = op(e2, e0)) & (e2 = op(e0, e2)) & (e2 = op(e2, e0))) | ~ sP37), inference(nnf_transformation, [], [f91])).
fof(f1123, plain, (~ spl66_93 | spl66_55), inference(avatar_split_clause, [], [f402, f819, f1118])).
fof(f402, plain, ((e2 = op(e0, e2)) | ~ sP37), inference(cnf_transformation, [], [f149])).
fof(f1122, plain, (~ spl66_93 | spl66_29), inference(avatar_split_clause, [], [f403, f709, f1118])).
fof(f403, plain, ((e0 = op(e2, e0)) | ~ sP37), inference(cnf_transformation, [], [f149])).
fof(f1121, plain, (~ spl66_93 | spl66_53), inference(avatar_split_clause, [], [f404, f811, f1118])).
fof(f404, plain, ((e0 = op(e0, e2)) | ~ sP37), inference(cnf_transformation, [], [f149])).
fof(f1116, plain, (~ spl66_92 | spl66_31), inference(avatar_split_clause, [], [f397, f717, f1110])).
fof(f397, plain, ((e2 = op(e2, e0)) | ~ sP38), inference(cnf_transformation, [], [f148])).
fof(f148, plain, (((e0 = op(e0, e3)) & (e0 = op(e3, e0)) & (e2 = op(e0, e2)) & (e2 = op(e2, e0))) | ~ sP38), inference(nnf_transformation, [], [f92])).
fof(f1115, plain, (~ spl66_92 | spl66_55), inference(avatar_split_clause, [], [f398, f819, f1110])).
fof(f398, plain, ((e2 = op(e0, e2)) | ~ sP38), inference(cnf_transformation, [], [f148])).
fof(f1114, plain, (~ spl66_92 | spl66_13), inference(avatar_split_clause, [], [f399, f641, f1110])).
fof(f399, plain, ((e0 = op(e3, e0)) | ~ sP38), inference(cnf_transformation, [], [f148])).
fof(f1113, plain, (~ spl66_92 | spl66_49), inference(avatar_split_clause, [], [f400, f794, f1110])).
fof(f400, plain, ((e0 = op(e0, e3)) | ~ sP38), inference(cnf_transformation, [], [f148])).
fof(f1108, plain, (~ spl66_91 | spl66_27), inference(avatar_split_clause, [], [f393, f700, f1102])).
fof(f393, plain, ((e2 = op(e2, e1)) | ~ sP39), inference(cnf_transformation, [], [f147])).
fof(f147, plain, (((e1 = op(e1, e0)) & (e1 = op(e0, e1)) & (e2 = op(e1, e2)) & (e2 = op(e2, e1))) | ~ sP39), inference(nnf_transformation, [], [f93])).
fof(f1107, plain, (~ spl66_91 | spl66_39), inference(avatar_split_clause, [], [f394, f751, f1102])).
fof(f394, plain, ((e2 = op(e1, e2)) | ~ sP39), inference(cnf_transformation, [], [f147])).
fof(f1106, plain, (~ spl66_91 | spl66_58), inference(avatar_split_clause, [], [f395, f832, f1102])).
fof(f395, plain, ((e1 = op(e0, e1)) | ~ sP39), inference(cnf_transformation, [], [f147])).
fof(f1105, plain, (~ spl66_91 | spl66_46), inference(avatar_split_clause, [], [f396, f781, f1102])).
fof(f396, plain, ((e1 = op(e1, e0)) | ~ sP39), inference(cnf_transformation, [], [f147])).
fof(f1100, plain, (~ spl66_90 | spl66_27), inference(avatar_split_clause, [], [f389, f700, f1094])).
fof(f389, plain, ((e2 = op(e2, e1)) | ~ sP40), inference(cnf_transformation, [], [f146])).
fof(f146, plain, (((e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e2 = op(e1, e2)) & (e2 = op(e2, e1))) | ~ sP40), inference(nnf_transformation, [], [f94])).
fof(f1099, plain, (~ spl66_90 | spl66_39), inference(avatar_split_clause, [], [f390, f751, f1094])).
fof(f390, plain, ((e2 = op(e1, e2)) | ~ sP40), inference(cnf_transformation, [], [f146])).
fof(f1098, plain, (~ spl66_90 | spl66_42), inference(avatar_split_clause, [], [f391, f764, f1094])).
fof(f391, plain, ((e1 = op(e1, e1)) | ~ sP40), inference(cnf_transformation, [], [f146])).
fof(f1092, plain, (~ spl66_89 | spl66_27), inference(avatar_split_clause, [], [f385, f700, f1086])).
fof(f385, plain, ((e2 = op(e2, e1)) | ~ sP41), inference(cnf_transformation, [], [f145])).
fof(f145, plain, (((e1 = op(e1, e2)) & (e1 = op(e2, e1)) & (e2 = op(e1, e2)) & (e2 = op(e2, e1))) | ~ sP41), inference(nnf_transformation, [], [f95])).
fof(f1091, plain, (~ spl66_89 | spl66_39), inference(avatar_split_clause, [], [f386, f751, f1086])).
fof(f386, plain, ((e2 = op(e1, e2)) | ~ sP41), inference(cnf_transformation, [], [f145])).
fof(f1090, plain, (~ spl66_89 | spl66_26), inference(avatar_split_clause, [], [f387, f696, f1086])).
fof(f387, plain, ((e1 = op(e2, e1)) | ~ sP41), inference(cnf_transformation, [], [f145])).
fof(f1089, plain, (~ spl66_89 | spl66_38), inference(avatar_split_clause, [], [f388, f747, f1086])).
fof(f388, plain, ((e1 = op(e1, e2)) | ~ sP41), inference(cnf_transformation, [], [f145])).
fof(f1084, plain, (~ spl66_88 | spl66_27), inference(avatar_split_clause, [], [f381, f700, f1078])).
fof(f381, plain, ((e2 = op(e2, e1)) | ~ sP42), inference(cnf_transformation, [], [f144])).
fof(f144, plain, (((e1 = op(e1, e3)) & (e1 = op(e3, e1)) & (e2 = op(e1, e2)) & (e2 = op(e2, e1))) | ~ sP42), inference(nnf_transformation, [], [f96])).
fof(f1083, plain, (~ spl66_88 | spl66_39), inference(avatar_split_clause, [], [f382, f751, f1078])).
fof(f382, plain, ((e2 = op(e1, e2)) | ~ sP42), inference(cnf_transformation, [], [f144])).
fof(f1082, plain, (~ spl66_88 | spl66_10), inference(avatar_split_clause, [], [f383, f628, f1078])).
fof(f383, plain, ((e1 = op(e3, e1)) | ~ sP42), inference(cnf_transformation, [], [f144])).
fof(f1081, plain, (~ spl66_88 | spl66_34), inference(avatar_split_clause, [], [f384, f730, f1078])).
fof(f384, plain, ((e1 = op(e1, e3)) | ~ sP42), inference(cnf_transformation, [], [f144])).
fof(f1076, plain, (~ spl66_87 | spl66_23), inference(avatar_split_clause, [], [f377, f683, f1070])).
fof(f377, plain, ((e2 = op(e2, e2)) | ~ sP43), inference(cnf_transformation, [], [f143])).
fof(f143, plain, (((e2 = op(e2, e0)) & (e2 = op(e0, e2)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP43), inference(nnf_transformation, [], [f97])).
fof(f1074, plain, (~ spl66_87 | spl66_55), inference(avatar_split_clause, [], [f379, f819, f1070])).
fof(f379, plain, ((e2 = op(e0, e2)) | ~ sP43), inference(cnf_transformation, [], [f143])).
fof(f1073, plain, (~ spl66_87 | spl66_31), inference(avatar_split_clause, [], [f380, f717, f1070])).
fof(f380, plain, ((e2 = op(e2, e0)) | ~ sP43), inference(cnf_transformation, [], [f143])).
fof(f1068, plain, (~ spl66_86 | spl66_23), inference(avatar_split_clause, [], [f373, f683, f1062])).
fof(f373, plain, ((e2 = op(e2, e2)) | ~ sP44), inference(cnf_transformation, [], [f142])).
fof(f142, plain, (((e2 = op(e2, e1)) & (e2 = op(e1, e2)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP44), inference(nnf_transformation, [], [f98])).
fof(f1065, plain, (~ spl66_86 | spl66_27), inference(avatar_split_clause, [], [f376, f700, f1062])).
fof(f376, plain, ((e2 = op(e2, e1)) | ~ sP44), inference(cnf_transformation, [], [f142])).
fof(f1060, plain, (~ spl66_85 | spl66_23), inference(avatar_split_clause, [], [f369, f683, f1054])).
fof(f369, plain, ((e2 = op(e2, e2)) | ~ sP45), inference(cnf_transformation, [], [f141])).
fof(f141, plain, (((e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP45), inference(nnf_transformation, [], [f99])).
fof(f1052, plain, (~ spl66_84 | spl66_23), inference(avatar_split_clause, [], [f365, f683, f1046])).
fof(f365, plain, ((e2 = op(e2, e2)) | ~ sP46), inference(cnf_transformation, [], [f140])).
fof(f140, plain, (((e2 = op(e2, e3)) & (e2 = op(e3, e2)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP46), inference(nnf_transformation, [], [f100])).
fof(f1049, plain, (~ spl66_84 | spl66_19), inference(avatar_split_clause, [], [f368, f666, f1046])).
fof(f368, plain, ((e2 = op(e2, e3)) | ~ sP46), inference(cnf_transformation, [], [f140])).
fof(f1044, plain, (~ spl66_83 | spl66_19), inference(avatar_split_clause, [], [f361, f666, f1038])).
fof(f361, plain, ((e2 = op(e2, e3)) | ~ sP47), inference(cnf_transformation, [], [f139])).
fof(f139, plain, (((e3 = op(e3, e0)) & (e3 = op(e0, e3)) & (e2 = op(e3, e2)) & (e2 = op(e2, e3))) | ~ sP47), inference(nnf_transformation, [], [f101])).
fof(f1043, plain, (~ spl66_83 | spl66_7), inference(avatar_split_clause, [], [f362, f615, f1038])).
fof(f362, plain, ((e2 = op(e3, e2)) | ~ sP47), inference(cnf_transformation, [], [f139])).
fof(f1042, plain, (~ spl66_83 | spl66_52), inference(avatar_split_clause, [], [f363, f806, f1038])).
fof(f363, plain, ((e3 = op(e0, e3)) | ~ sP47), inference(cnf_transformation, [], [f139])).
fof(f1041, plain, (~ spl66_83 | spl66_16), inference(avatar_split_clause, [], [f364, f653, f1038])).
fof(f364, plain, ((e3 = op(e3, e0)) | ~ sP47), inference(cnf_transformation, [], [f139])).
fof(f1036, plain, (~ spl66_82 | spl66_19), inference(avatar_split_clause, [], [f357, f666, f1030])).
fof(f357, plain, ((e2 = op(e2, e3)) | ~ sP48), inference(cnf_transformation, [], [f138])).
fof(f138, plain, (((e3 = op(e3, e1)) & (e3 = op(e1, e3)) & (e2 = op(e3, e2)) & (e2 = op(e2, e3))) | ~ sP48), inference(nnf_transformation, [], [f102])).
fof(f1035, plain, (~ spl66_82 | spl66_7), inference(avatar_split_clause, [], [f358, f615, f1030])).
fof(f358, plain, ((e2 = op(e3, e2)) | ~ sP48), inference(cnf_transformation, [], [f138])).
fof(f1034, plain, (~ spl66_82 | spl66_36), inference(avatar_split_clause, [], [f359, f738, f1030])).
fof(f359, plain, ((e3 = op(e1, e3)) | ~ sP48), inference(cnf_transformation, [], [f138])).
fof(f1033, plain, (~ spl66_82 | spl66_12), inference(avatar_split_clause, [], [f360, f636, f1030])).
fof(f360, plain, ((e3 = op(e3, e1)) | ~ sP48), inference(cnf_transformation, [], [f138])).
fof(f1028, plain, (~ spl66_81 | spl66_19), inference(avatar_split_clause, [], [f353, f666, f1022])).
fof(f353, plain, ((e2 = op(e2, e3)) | ~ sP49), inference(cnf_transformation, [], [f137])).
fof(f137, plain, (((e3 = op(e3, e2)) & (e3 = op(e2, e3)) & (e2 = op(e3, e2)) & (e2 = op(e2, e3))) | ~ sP49), inference(nnf_transformation, [], [f103])).
fof(f1026, plain, (~ spl66_81 | spl66_20), inference(avatar_split_clause, [], [f355, f670, f1022])).
fof(f355, plain, ((e3 = op(e2, e3)) | ~ sP49), inference(cnf_transformation, [], [f137])).
fof(f1025, plain, (~ spl66_81 | spl66_8), inference(avatar_split_clause, [], [f356, f619, f1022])).
fof(f356, plain, ((e3 = op(e3, e2)) | ~ sP49), inference(cnf_transformation, [], [f137])).
fof(f1018, plain, (~ spl66_80 | spl66_4), inference(avatar_split_clause, [], [f351, f602, f1014])).
fof(f351, plain, ((e3 = op(e3, e3)) | ~ sP50), inference(cnf_transformation, [], [f136])).
fof(f136, plain, (((e3 = op(e3, e3)) & (e3 = op(e3, e3)) & (e2 = op(e3, e2)) & (e2 = op(e2, e3))) | ~ sP50), inference(nnf_transformation, [], [f104])).
fof(f1010, plain, (~ spl66_79 | spl66_61), inference(avatar_split_clause, [], [f347, f845, f1006])).
fof(f347, plain, ((e0 = op(e0, e0)) | ~ sP51), inference(cnf_transformation, [], [f135])).
fof(f135, plain, (((e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e3 = op(e0, e3)) & (e3 = op(e3, e0))) | ~ sP51), inference(nnf_transformation, [], [f105])).
fof(f1004, plain, (~ spl66_78 | spl66_16), inference(avatar_split_clause, [], [f341, f653, f998])).
fof(f341, plain, ((e3 = op(e3, e0)) | ~ sP52), inference(cnf_transformation, [], [f134])).
fof(f134, plain, (((e0 = op(e0, e1)) & (e0 = op(e1, e0)) & (e3 = op(e0, e3)) & (e3 = op(e3, e0))) | ~ sP52), inference(nnf_transformation, [], [f106])).
fof(f1003, plain, (~ spl66_78 | spl66_52), inference(avatar_split_clause, [], [f342, f806, f998])).
fof(f342, plain, ((e3 = op(e0, e3)) | ~ sP52), inference(cnf_transformation, [], [f134])).
fof(f1002, plain, (~ spl66_78 | spl66_45), inference(avatar_split_clause, [], [f343, f777, f998])).
fof(f343, plain, ((e0 = op(e1, e0)) | ~ sP52), inference(cnf_transformation, [], [f134])).
fof(f1001, plain, (~ spl66_78 | spl66_57), inference(avatar_split_clause, [], [f344, f828, f998])).
fof(f344, plain, ((e0 = op(e0, e1)) | ~ sP52), inference(cnf_transformation, [], [f134])).
fof(f996, plain, (~ spl66_77 | spl66_16), inference(avatar_split_clause, [], [f337, f653, f990])).
fof(f337, plain, ((e3 = op(e3, e0)) | ~ sP53), inference(cnf_transformation, [], [f133])).
fof(f133, plain, (((e0 = op(e0, e2)) & (e0 = op(e2, e0)) & (e3 = op(e0, e3)) & (e3 = op(e3, e0))) | ~ sP53), inference(nnf_transformation, [], [f107])).
fof(f995, plain, (~ spl66_77 | spl66_52), inference(avatar_split_clause, [], [f338, f806, f990])).
fof(f338, plain, ((e3 = op(e0, e3)) | ~ sP53), inference(cnf_transformation, [], [f133])).
fof(f994, plain, (~ spl66_77 | spl66_29), inference(avatar_split_clause, [], [f339, f709, f990])).
fof(f339, plain, ((e0 = op(e2, e0)) | ~ sP53), inference(cnf_transformation, [], [f133])).
fof(f993, plain, (~ spl66_77 | spl66_53), inference(avatar_split_clause, [], [f340, f811, f990])).
fof(f340, plain, ((e0 = op(e0, e2)) | ~ sP53), inference(cnf_transformation, [], [f133])).
fof(f988, plain, (~ spl66_76 | spl66_16), inference(avatar_split_clause, [], [f333, f653, f982])).
fof(f333, plain, ((e3 = op(e3, e0)) | ~ sP54), inference(cnf_transformation, [], [f132])).
fof(f132, plain, (((e0 = op(e0, e3)) & (e0 = op(e3, e0)) & (e3 = op(e0, e3)) & (e3 = op(e3, e0))) | ~ sP54), inference(nnf_transformation, [], [f108])).
fof(f987, plain, (~ spl66_76 | spl66_52), inference(avatar_split_clause, [], [f334, f806, f982])).
fof(f334, plain, ((e3 = op(e0, e3)) | ~ sP54), inference(cnf_transformation, [], [f132])).
fof(f986, plain, (~ spl66_76 | spl66_13), inference(avatar_split_clause, [], [f335, f641, f982])).
fof(f335, plain, ((e0 = op(e3, e0)) | ~ sP54), inference(cnf_transformation, [], [f132])).
fof(f985, plain, (~ spl66_76 | spl66_49), inference(avatar_split_clause, [], [f336, f794, f982])).
fof(f336, plain, ((e0 = op(e0, e3)) | ~ sP54), inference(cnf_transformation, [], [f132])).
fof(f980, plain, (~ spl66_75 | spl66_12), inference(avatar_split_clause, [], [f329, f636, f974])).
fof(f329, plain, ((e3 = op(e3, e1)) | ~ sP55), inference(cnf_transformation, [], [f131])).
fof(f131, plain, (((e1 = op(e1, e0)) & (e1 = op(e0, e1)) & (e3 = op(e1, e3)) & (e3 = op(e3, e1))) | ~ sP55), inference(nnf_transformation, [], [f109])).
fof(f979, plain, (~ spl66_75 | spl66_36), inference(avatar_split_clause, [], [f330, f738, f974])).
fof(f330, plain, ((e3 = op(e1, e3)) | ~ sP55), inference(cnf_transformation, [], [f131])).
fof(f978, plain, (~ spl66_75 | spl66_58), inference(avatar_split_clause, [], [f331, f832, f974])).
fof(f331, plain, ((e1 = op(e0, e1)) | ~ sP55), inference(cnf_transformation, [], [f131])).
fof(f977, plain, (~ spl66_75 | spl66_46), inference(avatar_split_clause, [], [f332, f781, f974])).
fof(f332, plain, ((e1 = op(e1, e0)) | ~ sP55), inference(cnf_transformation, [], [f131])).
fof(f972, plain, (~ spl66_74 | spl66_12), inference(avatar_split_clause, [], [f325, f636, f966])).
fof(f325, plain, ((e3 = op(e3, e1)) | ~ sP56), inference(cnf_transformation, [], [f130])).
fof(f130, plain, (((e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e3 = op(e1, e3)) & (e3 = op(e3, e1))) | ~ sP56), inference(nnf_transformation, [], [f110])).
fof(f971, plain, (~ spl66_74 | spl66_36), inference(avatar_split_clause, [], [f326, f738, f966])).
fof(f326, plain, ((e3 = op(e1, e3)) | ~ sP56), inference(cnf_transformation, [], [f130])).
fof(f970, plain, (~ spl66_74 | spl66_42), inference(avatar_split_clause, [], [f327, f764, f966])).
fof(f327, plain, ((e1 = op(e1, e1)) | ~ sP56), inference(cnf_transformation, [], [f130])).
fof(f964, plain, (~ spl66_73 | spl66_12), inference(avatar_split_clause, [], [f321, f636, f958])).
fof(f321, plain, ((e3 = op(e3, e1)) | ~ sP57), inference(cnf_transformation, [], [f129])).
fof(f129, plain, (((e1 = op(e1, e2)) & (e1 = op(e2, e1)) & (e3 = op(e1, e3)) & (e3 = op(e3, e1))) | ~ sP57), inference(nnf_transformation, [], [f111])).
fof(f963, plain, (~ spl66_73 | spl66_36), inference(avatar_split_clause, [], [f322, f738, f958])).
fof(f322, plain, ((e3 = op(e1, e3)) | ~ sP57), inference(cnf_transformation, [], [f129])).
fof(f962, plain, (~ spl66_73 | spl66_26), inference(avatar_split_clause, [], [f323, f696, f958])).
fof(f323, plain, ((e1 = op(e2, e1)) | ~ sP57), inference(cnf_transformation, [], [f129])).
fof(f961, plain, (~ spl66_73 | spl66_38), inference(avatar_split_clause, [], [f324, f747, f958])).
fof(f324, plain, ((e1 = op(e1, e2)) | ~ sP57), inference(cnf_transformation, [], [f129])).
fof(f956, plain, (~ spl66_72 | spl66_12), inference(avatar_split_clause, [], [f317, f636, f950])).
fof(f317, plain, ((e3 = op(e3, e1)) | ~ sP58), inference(cnf_transformation, [], [f128])).
fof(f128, plain, (((e1 = op(e1, e3)) & (e1 = op(e3, e1)) & (e3 = op(e1, e3)) & (e3 = op(e3, e1))) | ~ sP58), inference(nnf_transformation, [], [f112])).
fof(f955, plain, (~ spl66_72 | spl66_36), inference(avatar_split_clause, [], [f318, f738, f950])).
fof(f318, plain, ((e3 = op(e1, e3)) | ~ sP58), inference(cnf_transformation, [], [f128])).
fof(f954, plain, (~ spl66_72 | spl66_10), inference(avatar_split_clause, [], [f319, f628, f950])).
fof(f319, plain, ((e1 = op(e3, e1)) | ~ sP58), inference(cnf_transformation, [], [f128])).
fof(f953, plain, (~ spl66_72 | spl66_34), inference(avatar_split_clause, [], [f320, f730, f950])).
fof(f320, plain, ((e1 = op(e1, e3)) | ~ sP58), inference(cnf_transformation, [], [f128])).
fof(f948, plain, (~ spl66_71 | spl66_8), inference(avatar_split_clause, [], [f313, f619, f942])).
fof(f313, plain, ((e3 = op(e3, e2)) | ~ sP59), inference(cnf_transformation, [], [f127])).
fof(f127, plain, (((e2 = op(e2, e0)) & (e2 = op(e0, e2)) & (e3 = op(e2, e3)) & (e3 = op(e3, e2))) | ~ sP59), inference(nnf_transformation, [], [f113])).
fof(f947, plain, (~ spl66_71 | spl66_20), inference(avatar_split_clause, [], [f314, f670, f942])).
fof(f314, plain, ((e3 = op(e2, e3)) | ~ sP59), inference(cnf_transformation, [], [f127])).
fof(f946, plain, (~ spl66_71 | spl66_55), inference(avatar_split_clause, [], [f315, f819, f942])).
fof(f315, plain, ((e2 = op(e0, e2)) | ~ sP59), inference(cnf_transformation, [], [f127])).
fof(f945, plain, (~ spl66_71 | spl66_31), inference(avatar_split_clause, [], [f316, f717, f942])).
fof(f316, plain, ((e2 = op(e2, e0)) | ~ sP59), inference(cnf_transformation, [], [f127])).
fof(f940, plain, (~ spl66_70 | spl66_8), inference(avatar_split_clause, [], [f309, f619, f934])).
fof(f309, plain, ((e3 = op(e3, e2)) | ~ sP60), inference(cnf_transformation, [], [f126])).
fof(f126, plain, (((e2 = op(e2, e1)) & (e2 = op(e1, e2)) & (e3 = op(e2, e3)) & (e3 = op(e3, e2))) | ~ sP60), inference(nnf_transformation, [], [f114])).
fof(f939, plain, (~ spl66_70 | spl66_20), inference(avatar_split_clause, [], [f310, f670, f934])).
fof(f310, plain, ((e3 = op(e2, e3)) | ~ sP60), inference(cnf_transformation, [], [f126])).
fof(f938, plain, (~ spl66_70 | spl66_39), inference(avatar_split_clause, [], [f311, f751, f934])).
fof(f311, plain, ((e2 = op(e1, e2)) | ~ sP60), inference(cnf_transformation, [], [f126])).
fof(f937, plain, (~ spl66_70 | spl66_27), inference(avatar_split_clause, [], [f312, f700, f934])).
fof(f312, plain, ((e2 = op(e2, e1)) | ~ sP60), inference(cnf_transformation, [], [f126])).
fof(f931, plain, (~ spl66_69 | spl66_20), inference(avatar_split_clause, [], [f306, f670, f926])).
fof(f306, plain, ((e3 = op(e2, e3)) | ~ sP61), inference(cnf_transformation, [], [f125])).
fof(f125, plain, (((e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e3 = op(e2, e3)) & (e3 = op(e3, e2))) | ~ sP61), inference(nnf_transformation, [], [f115])).
fof(f930, plain, (~ spl66_69 | spl66_23), inference(avatar_split_clause, [], [f307, f683, f926])).
fof(f307, plain, ((e2 = op(e2, e2)) | ~ sP61), inference(cnf_transformation, [], [f125])).
fof(f924, plain, (~ spl66_68 | spl66_8), inference(avatar_split_clause, [], [f301, f619, f918])).
fof(f301, plain, ((e3 = op(e3, e2)) | ~ sP62), inference(cnf_transformation, [], [f124])).
fof(f124, plain, (((e2 = op(e2, e3)) & (e2 = op(e3, e2)) & (e3 = op(e2, e3)) & (e3 = op(e3, e2))) | ~ sP62), inference(nnf_transformation, [], [f116])).
fof(f923, plain, (~ spl66_68 | spl66_20), inference(avatar_split_clause, [], [f302, f670, f918])).
fof(f302, plain, ((e3 = op(e2, e3)) | ~ sP62), inference(cnf_transformation, [], [f124])).
fof(f921, plain, (~ spl66_68 | spl66_19), inference(avatar_split_clause, [], [f304, f666, f918])).
fof(f304, plain, ((e2 = op(e2, e3)) | ~ sP62), inference(cnf_transformation, [], [f124])).
fof(f916, plain, (~ spl66_67 | spl66_4), inference(avatar_split_clause, [], [f297, f602, f910])).
fof(f297, plain, ((e3 = op(e3, e3)) | ~ sP63), inference(cnf_transformation, [], [f123])).
fof(f123, plain, (((e3 = op(e3, e0)) & (e3 = op(e0, e3)) & (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ~ sP63), inference(nnf_transformation, [], [f117])).
fof(f908, plain, (~ spl66_66 | spl66_4), inference(avatar_split_clause, [], [f293, f602, f902])).
fof(f293, plain, ((e3 = op(e3, e3)) | ~ sP64), inference(cnf_transformation, [], [f122])).
fof(f122, plain, (((e3 = op(e3, e1)) & (e3 = op(e1, e3)) & (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ~ sP64), inference(nnf_transformation, [], [f118])).
fof(f906, plain, (~ spl66_66 | spl66_36), inference(avatar_split_clause, [], [f295, f738, f902])).
fof(f295, plain, ((e3 = op(e1, e3)) | ~ sP64), inference(cnf_transformation, [], [f122])).
fof(f900, plain, (~ spl66_65 | spl66_4), inference(avatar_split_clause, [], [f289, f602, f894])).
fof(f289, plain, ((e3 = op(e3, e3)) | ~ sP65), inference(cnf_transformation, [], [f121])).
fof(f121, plain, (((e3 = op(e3, e2)) & (e3 = op(e2, e3)) & (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ~ sP65), inference(nnf_transformation, [], [f119])).
fof(f888, plain, (spl66_63 | spl66_59 | spl66_55 | spl66_51), inference(avatar_split_clause, [], [f207, f802, f819, f836, f853])).
fof(f207, plain, ((e2 = op(e0, e3)) | (e2 = op(e0, e2)) | (e2 = op(e0, e1)) | (op(e0, e0) = e2)), inference(cnf_transformation, [], [f2])).
fof(f2, plain, (((e3 = op(e3, e3)) | (e3 = op(e2, e3)) | (e3 = op(e1, e3)) | (e3 = op(e0, e3))) & ((e3 = op(e3, e3)) | (e3 = op(e3, e2)) | (e3 = op(e3, e1)) | (e3 = op(e3, e0))) & ((e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(e0, e3))) & ((e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, e0))) & ((e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(e0, e3))) & ((e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, e0))) & ((e0 = op(e3, e3)) | (e0 = op(e2, e3)) | (e0 = op(e1, e3)) | (e0 = op(e0, e3))) & ((e0 = op(e3, e3)) | (e0 = op(e3, e2)) | (e0 = op(e3, e1)) | (e0 = op(e3, e0))) & ((e3 = op(e3, e2)) | (e3 = op(e2, e2)) | (e3 = op(e1, e2)) | (e3 = op(e0, e2))) & ((e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, e0))) & ((e2 = op(e3, e2)) | (e2 = op(e2, e2)) | (e2 = op(e1, e2)) | (e2 = op(e0, e2))) & ((e2 = op(e2, e3)) | (e2 = op(e2, e2)) | (e2 = op(e2, e1)) | (e2 = op(e2, e0))) & ((e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(e0, e2))) & ((e1 = op(e2, e3)) | (e1 = op(e2, e2)) | (e1 = op(e2, e1)) | (e1 = op(e2, e0))) & ((e0 = op(e3, e2)) | (e0 = op(e2, e2)) | (e0 = op(e1, e2)) | (e0 = op(e0, e2))) & ((e0 = op(e2, e3)) | (e0 = op(e2, e2)) | (e0 = op(e2, e1)) | (e0 = op(e2, e0))) & ((e3 = op(e3, e1)) | (e3 = op(e2, e1)) | (e3 = op(e1, e1)) | (e3 = op(e0, e1))) & ((e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, e0))) & ((e2 = op(e3, e1)) | (e2 = op(e2, e1)) | (e2 = op(e1, e1)) | (e2 = op(e0, e1))) & ((e2 = op(e1, e3)) | (e2 = op(e1, e2)) | (e2 = op(e1, e1)) | (e2 = op(e1, e0))) & ((e1 = op(e3, e1)) | (e1 = op(e2, e1)) | (e1 = op(e1, e1)) | (e1 = op(e0, e1))) & ((e1 = op(e1, e3)) | (e1 = op(e1, e2)) | (e1 = op(e1, e1)) | (e1 = op(e1, e0))) & ((e0 = op(e3, e1)) | (e0 = op(e2, e1)) | (e0 = op(e1, e1)) | (e0 = op(e0, e1))) & ((e0 = op(e1, e3)) | (e0 = op(e1, e2)) | (e0 = op(e1, e1)) | (e0 = op(e1, e0))) & ((e3 = op(e3, e0)) | (e3 = op(e2, e0)) | (e3 = op(e1, e0)) | (op(e0, e0) = e3)) & ((e3 = op(e0, e3)) | (e3 = op(e0, e2)) | (e3 = op(e0, e1)) | (op(e0, e0) = e3)) & ((e2 = op(e3, e0)) | (e2 = op(e2, e0)) | (e2 = op(e1, e0)) | (op(e0, e0) = e2)) & ((e2 = op(e0, e3)) | (e2 = op(e0, e2)) | (e2 = op(e0, e1)) | (op(e0, e0) = e2)) & ((e1 = op(e3, e0)) | (e1 = op(e2, e0)) | (e1 = op(e1, e0)) | (op(e0, e0) = e1)) & ((e1 = op(e0, e3)) | (e1 = op(e0, e2)) | (e1 = op(e0, e1)) | (op(e0, e0) = e1)) & ((e0 = op(e3, e0)) | (e0 = op(e2, e0)) | (e0 = op(e1, e0)) | (e0 = op(e0, e0))) & ((e0 = op(e0, e3)) | (e0 = op(e0, e2)) | (e0 = op(e0, e1)) | (e0 = op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG143+1.p', ax2)).
fof(f887, plain, (spl66_63 | spl66_47 | spl66_31 | spl66_15), inference(avatar_split_clause, [], [f208, f649, f717, f785, f853])).
fof(f208, plain, ((e2 = op(e3, e0)) | (e2 = op(e2, e0)) | (e2 = op(e1, e0)) | (op(e0, e0) = e2)), inference(cnf_transformation, [], [f2])).
fof(f885, plain, (spl66_64 | spl66_48 | spl66_32 | spl66_16), inference(avatar_split_clause, [], [f210, f653, f721, f789, f857])).
fof(f210, plain, ((e3 = op(e3, e0)) | (e3 = op(e2, e0)) | (e3 = op(e1, e0)) | (op(e0, e0) = e3)), inference(cnf_transformation, [], [f2])).
fof(f884, plain, (spl66_45 | spl66_41 | spl66_37 | spl66_33), inference(avatar_split_clause, [], [f211, f726, f743, f760, f777])).
fof(f211, plain, ((e0 = op(e1, e3)) | (e0 = op(e1, e2)) | (e0 = op(e1, e1)) | (e0 = op(e1, e0))), inference(cnf_transformation, [], [f2])).
fof(f883, plain, (spl66_57 | spl66_41 | spl66_25 | spl66_9), inference(avatar_split_clause, [], [f212, f624, f692, f760, f828])).
fof(f212, plain, ((e0 = op(e3, e1)) | (e0 = op(e2, e1)) | (e0 = op(e1, e1)) | (e0 = op(e0, e1))), inference(cnf_transformation, [], [f2])).
fof(f882, plain, (spl66_46 | spl66_42 | spl66_38 | spl66_34), inference(avatar_split_clause, [], [f213, f730, f747, f764, f781])).
fof(f213, plain, ((e1 = op(e1, e3)) | (e1 = op(e1, e2)) | (e1 = op(e1, e1)) | (e1 = op(e1, e0))), inference(cnf_transformation, [], [f2])).
fof(f880, plain, (spl66_47 | spl66_43 | spl66_39 | spl66_35), inference(avatar_split_clause, [], [f215, f734, f751, f768, f785])).
fof(f215, plain, ((e2 = op(e1, e3)) | (e2 = op(e1, e2)) | (e2 = op(e1, e1)) | (e2 = op(e1, e0))), inference(cnf_transformation, [], [f2])).
fof(f879, plain, (spl66_59 | spl66_43 | spl66_27 | spl66_11), inference(avatar_split_clause, [], [f216, f632, f700, f768, f836])).
fof(f216, plain, ((e2 = op(e3, e1)) | (e2 = op(e2, e1)) | (e2 = op(e1, e1)) | (e2 = op(e0, e1))), inference(cnf_transformation, [], [f2])).
fof(f878, plain, (spl66_48 | spl66_44 | spl66_40 | spl66_36), inference(avatar_split_clause, [], [f217, f738, f755, f772, f789])).
fof(f217, plain, ((e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, e0))), inference(cnf_transformation, [], [f2])).
fof(f876, plain, (spl66_29 | spl66_25 | spl66_21 | spl66_17), inference(avatar_split_clause, [], [f219, f658, f675, f692, f709])).
fof(f219, plain, ((e0 = op(e2, e3)) | (e0 = op(e2, e2)) | (e0 = op(e2, e1)) | (e0 = op(e2, e0))), inference(cnf_transformation, [], [f2])).
fof(f875, plain, (spl66_53 | spl66_37 | spl66_21 | spl66_5), inference(avatar_split_clause, [], [f220, f607, f675, f743, f811])).
fof(f220, plain, ((e0 = op(e3, e2)) | (e0 = op(e2, e2)) | (e0 = op(e1, e2)) | (e0 = op(e0, e2))), inference(cnf_transformation, [], [f2])).
fof(f873, plain, (spl66_54 | spl66_38 | spl66_22 | spl66_6), inference(avatar_split_clause, [], [f222, f611, f679, f747, f815])).
fof(f222, plain, ((e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(e0, e2))), inference(cnf_transformation, [], [f2])).
fof(f870, plain, (spl66_32 | spl66_28 | spl66_24 | spl66_20), inference(avatar_split_clause, [], [f225, f670, f687, f704, f721])).
fof(f225, plain, ((e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, e0))), inference(cnf_transformation, [], [f2])).
fof(f868, plain, (spl66_13 | spl66_9 | spl66_5 | spl66_1), inference(avatar_split_clause, [], [f227, f590, f607, f624, f641])).
fof(f227, plain, ((e0 = op(e3, e3)) | (e0 = op(e3, e2)) | (e0 = op(e3, e1)) | (e0 = op(e3, e0))), inference(cnf_transformation, [], [f2])).
fof(f866, plain, (spl66_14 | spl66_10 | spl66_6 | spl66_2), inference(avatar_split_clause, [], [f229, f594, f611, f628, f645])).
fof(f229, plain, ((e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, e0))), inference(cnf_transformation, [], [f2])).
fof(f865, plain, (spl66_50 | spl66_34 | spl66_18 | spl66_2), inference(avatar_split_clause, [], [f230, f594, f662, f730, f798])).
fof(f230, plain, ((e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f864, plain, (spl66_15 | spl66_11 | spl66_7 | spl66_3), inference(avatar_split_clause, [], [f231, f598, f615, f632, f649])).
fof(f231, plain, ((e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, e0))), inference(cnf_transformation, [], [f2])).
fof(f863, plain, (spl66_51 | spl66_35 | spl66_19 | spl66_3), inference(avatar_split_clause, [], [f232, f598, f666, f734, f802])).
fof(f232, plain, ((e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f861, plain, (spl66_52 | spl66_36 | spl66_20 | spl66_4), inference(avatar_split_clause, [], [f234, f602, f670, f738, f806])).
fof(f234, plain, ((e3 = op(e3, e3)) | (e3 = op(e2, e3)) | (e3 = op(e1, e3)) | (e3 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f860, plain, (spl66_61 | spl66_62 | spl66_63 | spl66_64), inference(avatar_split_clause, [], [f187, f857, f853, f849, f845])).
fof(f187, plain, ((op(e0, e0) = e3) | (op(e0, e0) = e2) | (op(e0, e0) = e1) | (e0 = op(e0, e0))), inference(cnf_transformation, [], [f1])).
fof(f1, plain, (((e3 = op(e3, e3)) | (e2 = op(e3, e3)) | (e1 = op(e3, e3)) | (e0 = op(e3, e3))) & ((e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (e0 = op(e3, e2))) & ((e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (e0 = op(e3, e1))) & ((e3 = op(e3, e0)) | (e2 = op(e3, e0)) | (e1 = op(e3, e0)) | (e0 = op(e3, e0))) & ((e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))) & ((e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (e0 = op(e2, e2))) & ((e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))) & ((e3 = op(e2, e0)) | (e2 = op(e2, e0)) | (e1 = op(e2, e0)) | (e0 = op(e2, e0))) & ((e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (e0 = op(e1, e3))) & ((e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (e0 = op(e1, e2))) & ((e3 = op(e1, e1)) | (e2 = op(e1, e1)) | (e1 = op(e1, e1)) | (e0 = op(e1, e1))) & ((e3 = op(e1, e0)) | (e2 = op(e1, e0)) | (e1 = op(e1, e0)) | (e0 = op(e1, e0))) & ((e3 = op(e0, e3)) | (e2 = op(e0, e3)) | (e1 = op(e0, e3)) | (e0 = op(e0, e3))) & ((e3 = op(e0, e2)) | (e2 = op(e0, e2)) | (e1 = op(e0, e2)) | (e0 = op(e0, e2))) & ((e3 = op(e0, e1)) | (e2 = op(e0, e1)) | (e1 = op(e0, e1)) | (e0 = op(e0, e1))) & ((op(e0, e0) = e3) | (op(e0, e0) = e2) | (op(e0, e0) = e1) | (e0 = op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG143+1.p', ax1)).
fof(f843, plain, (spl66_57 | spl66_58 | spl66_59 | spl66_60), inference(avatar_split_clause, [], [f188, f840, f836, f832, f828])).
fof(f188, plain, ((e3 = op(e0, e1)) | (e2 = op(e0, e1)) | (e1 = op(e0, e1)) | (e0 = op(e0, e1))), inference(cnf_transformation, [], [f1])).
fof(f826, plain, (spl66_53 | spl66_54 | spl66_55 | spl66_56), inference(avatar_split_clause, [], [f189, f823, f819, f815, f811])).
fof(f189, plain, ((e3 = op(e0, e2)) | (e2 = op(e0, e2)) | (e1 = op(e0, e2)) | (e0 = op(e0, e2))), inference(cnf_transformation, [], [f1])).
fof(f809, plain, (spl66_49 | spl66_50 | spl66_51 | spl66_52), inference(avatar_split_clause, [], [f190, f806, f802, f798, f794])).
fof(f190, plain, ((e3 = op(e0, e3)) | (e2 = op(e0, e3)) | (e1 = op(e0, e3)) | (e0 = op(e0, e3))), inference(cnf_transformation, [], [f1])).
fof(f792, plain, (spl66_45 | spl66_46 | spl66_47 | spl66_48), inference(avatar_split_clause, [], [f191, f789, f785, f781, f777])).
fof(f191, plain, ((e3 = op(e1, e0)) | (e2 = op(e1, e0)) | (e1 = op(e1, e0)) | (e0 = op(e1, e0))), inference(cnf_transformation, [], [f1])).
fof(f775, plain, (spl66_41 | spl66_42 | spl66_43 | spl66_44), inference(avatar_split_clause, [], [f192, f772, f768, f764, f760])).
fof(f192, plain, ((e3 = op(e1, e1)) | (e2 = op(e1, e1)) | (e1 = op(e1, e1)) | (e0 = op(e1, e1))), inference(cnf_transformation, [], [f1])).
fof(f758, plain, (spl66_37 | spl66_38 | spl66_39 | spl66_40), inference(avatar_split_clause, [], [f193, f755, f751, f747, f743])).
fof(f193, plain, ((e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (e0 = op(e1, e2))), inference(cnf_transformation, [], [f1])).
fof(f741, plain, (spl66_33 | spl66_34 | spl66_35 | spl66_36), inference(avatar_split_clause, [], [f194, f738, f734, f730, f726])).
fof(f194, plain, ((e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (e0 = op(e1, e3))), inference(cnf_transformation, [], [f1])).
fof(f724, plain, (spl66_29 | spl66_30 | spl66_31 | spl66_32), inference(avatar_split_clause, [], [f195, f721, f717, f713, f709])).
fof(f195, plain, ((e3 = op(e2, e0)) | (e2 = op(e2, e0)) | (e1 = op(e2, e0)) | (e0 = op(e2, e0))), inference(cnf_transformation, [], [f1])).
fof(f707, plain, (spl66_25 | spl66_26 | spl66_27 | spl66_28), inference(avatar_split_clause, [], [f196, f704, f700, f696, f692])).
fof(f196, plain, ((e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))), inference(cnf_transformation, [], [f1])).
fof(f690, plain, (spl66_21 | spl66_22 | spl66_23 | spl66_24), inference(avatar_split_clause, [], [f197, f687, f683, f679, f675])).
fof(f197, plain, ((e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (e0 = op(e2, e2))), inference(cnf_transformation, [], [f1])).
fof(f673, plain, (spl66_17 | spl66_18 | spl66_19 | spl66_20), inference(avatar_split_clause, [], [f198, f670, f666, f662, f658])).
fof(f198, plain, ((e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))), inference(cnf_transformation, [], [f1])).
fof(f656, plain, (spl66_13 | spl66_14 | spl66_15 | spl66_16), inference(avatar_split_clause, [], [f199, f653, f649, f645, f641])).
fof(f199, plain, ((e3 = op(e3, e0)) | (e2 = op(e3, e0)) | (e1 = op(e3, e0)) | (e0 = op(e3, e0))), inference(cnf_transformation, [], [f1])).
fof(f639, plain, (spl66_9 | spl66_10 | spl66_11 | spl66_12), inference(avatar_split_clause, [], [f200, f636, f632, f628, f624])).
fof(f200, plain, ((e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (e0 = op(e3, e1))), inference(cnf_transformation, [], [f1])).
fof(f622, plain, (spl66_5 | spl66_6 | spl66_7 | spl66_8), inference(avatar_split_clause, [], [f201, f619, f615, f611, f607])).
fof(f201, plain, ((e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (e0 = op(e3, e2))), inference(cnf_transformation, [], [f1])).
fof(f605, plain, (spl66_1 | spl66_2 | spl66_3 | spl66_4), inference(avatar_split_clause, [], [f202, f602, f598, f594, f590])).
fof(f202, plain, ((e3 = op(e3, e3)) | (e2 = op(e3, e3)) | (e1 = op(e3, e3)) | (e0 = op(e3, e3))), inference(cnf_transformation, [], [f1])).