fof(f4141, plain, $false, inference(avatar_sat_refutation, [], [f617, f634, f651, f668, f685, f702, f719, f736, f753, f770, f787, f804, f821, f838, f855, f872, f873, f874, f875, f876, f877, f878, f879, f880, f881, f882, f884, f885, f886, f887, f888, f889, f890, f891, f892, f894, f895, f896, f897, f898, f899, f900, f901, f902, f903, f904, f909, f910, f911, f921, f926, f931, f936, f937, f938, f943, f948, f953, f958, f963, f964, f965, f970, f975, f980, f985, f992, f993, f1000, f1001, f1008, f1009, f1015, f1016, f1022, f1023, f1030, f1031, f1032, f1033, f1038, f1039, f1040, f1041, f1047, f1048, f1049, f1054, f1055, f1056, f1057, f1062, f1063, f1070, f1071, f1072, f1073, f1078, f1079, f1080, f1081, f1086, f1087, f1088, f1089, f1094, f1095, f1096, f1097, f1102, f1103, f1104, f1105, f1110, f1111, f1119, f1120, f1126, f1127, f1128, f1129, f1134, f1135, f1136, f1137, f1143, f1144, f1145, f1152, f1153, f1160, f1161, f1168, f1169, f1174, f1175, f1176, f1177, f1183, f1184, f1190, f1191, f1198, f1199, f1200, f1201, f1206, f1207, f1208, f1209, f1215, f1216, f1222, f1223, f1224, f1225, f1230, f1231, f1238, f1239, f1241, f1246, f1247, f1248, f1249, f1255, f1256, f1262, f1263, f1264, f1265, f1270, f1271, f1272, f1273, f1278, f1279, f1280, f1281, f1287, f1288, f1294, f1295, f1296, f1297, f1304, f1305, f1312, f1313, f1320, f1321, f1328, f1329, f1334, f1335, f1336, f1337, f1342, f1343, f1344, f1345, f1350, f1351, f1353, f1358, f1359, f1366, f1367, f1374, f1375, f1376, f1377, f1382, f1383, f1384, f1385, f1391, f1392, f1398, f1399, f1400, f1401, f1406, f1407, f1409, f1414, f1415, f1416, f1417, f1423, f1424, f1430, f1431, f1432, f1433, f1438, f1439, f1440, f1441, f1446, f1447, f1455, f1456, f1464, f1465, f1472, f1473, f1480, f1481, f1488, f1489, f1490, f1491, f1492, f1497, f1502, f1507, f1512, f1515, f1516, f1525, f1530, f1539, f1540, f1545, f1546, f1555, f1560, f1569, f1570, f1575, f1576, f1585, f1590, f1599, f1600, f1605, f1606, f1615, f1620, f1629, f1630, f1635, f1636, f1646, f1648, f1650, f1651, f1656, f1678, f1701, f1702, f1703, f1708, f1717, f1727, f1739, f1748, f1752, f1755, f1764, f1767, f1775, f1783, f1803, f1817, f1821, f1841, f1859, f1863, f1870, f1873, f1880, f1893, f1904, f1913, f1918, f1932, f1942, f1948, f1961, f1975, f1991, f2000, f2002, f2013, f2017, f2024, f2028, f2033, f2036, f2040, f2047, f2053, f2070, f2082, f2086, f2089, f2121, f2126, f2145, f2157, f2161, f2165, f2171, f2174, f2176, f2179, f2189, f2202, f2205, f2216, f2225, f2226, f2240, f2244, f2251, f2267, f2270, f2272, f2276, f2298, f2305, f2314, f2316, f2324, f2327, f2337, f2355, f2393, f2416, f2420, f2434, f2435, f2457, f2468, f2473, f2476, f2480, f2488, f2515, f2520, f2526, f2541, f2545, f2548, f2551, f2572, f2589, f2602, f2623, f2627, f2642, f2650, f2676, f2690, f2692, f2705, f2706, f2720, f2738, f2750, f2758, f2760, f2765, f2790, f2795, f2799, f2800, f2803, f2814, f2827, f2828, f2840, f2844, f2858, f2863, f2879, f2885, f2896, f2920, f2921, f2923, f2937, f2945, f2946, f2951, f2953, f2960, f2961, f2973, f2975, f2983, f2994, f3022, f3025, f3032, f3042, f3043, f3064, f3066, f3079, f3086, f3105, f3106, f3120, f3127, f3132, f3139, f3140, f3146, f3157, f3176, f3184, f3188, f3221, f3229, f3240, f3245, f3256, f3257, f3268, f3271, f3287, f3288, f3298, f3299, f3326, f3329, f3337, f3339, f3341, f3347, f3349, f3368, f3370, f3387, f3395, f3416, f3417, f3420, f3428, f3435, f3455, f3462, f3463, f3474, f3482, f3494, f3512, f3526, f3527, f3543, f3544, f3556, f3577, f3581, f3604, f3607, f3615, f3617, f3620, f3633, f3649, f3650, f3651, f3661, f3662, f3663, f3668, f3675, f3687, f3688, f3704, f3714, f3720, f3721, f3731, f3736, f3743, f3745, f3748, f3753, f3754, f3770, f3782, f3783, f3790, f3798, f3802, f3813, f3841, f3849, f3869, f3870, f3880, f3881, f3882, f3883, f3894, f3906, f3920, f3925, f3932, f3945, f3949, f3955, f3956, f3968, f3976, f3989, f3995, f3999, f4005, f4010, f4011, f4025, f4037, f4039, f4041, f4049, f4055, f4061, f4062, f4068, f4074, f4075, f4081, f4086, f4096, f4112, f4114, f4125, f4129])).
fof(f4129, plain, (~ spl66_4 | ~ spl66_12), inference(avatar_split_clause, [], [f4127, f648, f614])).
fof(f614, plain, (spl66_4 <=> (e3 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_4])])).
fof(f648, plain, (spl66_12 <=> (e3 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_12])])).
fof(f4127, plain, (~ (e3 = op(e3, e3)) | ~ spl66_12), inference(backward_demodulation, [], [f281, f650])).
fof(f650, plain, ((e3 = op(e3, e1)) | ~ spl66_12), inference(avatar_component_clause, [], [f648])).
fof(f281, plain, ~ (op(e3, e1) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f3, plain, (~ (op(e3, e2) = op(e3, e3)) & ~ (op(e3, e1) = op(e3, e3)) & ~ (op(e3, e1) = op(e3, e2)) & ~ (op(e3, e0) = op(e3, e3)) & ~ (op(e3, e0) = op(e3, e2)) & ~ (op(e3, e0) = op(e3, e1)) & ~ (op(e2, e2) = op(e2, e3)) & ~ (op(e2, e1) = op(e2, e3)) & ~ (op(e2, e1) = op(e2, e2)) & ~ (op(e2, e0) = op(e2, e3)) & ~ (op(e2, e0) = op(e2, e2)) & ~ (op(e2, e0) = op(e2, e1)) & ~ (op(e1, e2) = op(e1, e3)) & ~ (op(e1, e1) = op(e1, e3)) & ~ (op(e1, e1) = op(e1, e2)) & ~ (op(e1, e0) = op(e1, e3)) & ~ (op(e1, e0) = op(e1, e2)) & ~ (op(e1, e0) = op(e1, e1)) & ~ (op(e0, e2) = op(e0, e3)) & ~ (op(e0, e1) = op(e0, e3)) & ~ (op(e0, e1) = op(e0, e2)) & ~ (op(e0, e0) = op(e0, e3)) & ~ (op(e0, e0) = op(e0, e2)) & ~ (op(e0, e0) = op(e0, e1)) & ~ (op(e2, e3) = op(e3, e3)) & ~ (op(e1, e3) = op(e3, e3)) & ~ (op(e1, e3) = op(e2, e3)) & ~ (op(e0, e3) = op(e3, e3)) & ~ (op(e0, e3) = op(e2, e3)) & ~ (op(e0, e3) = op(e1, e3)) & ~ (op(e2, e2) = op(e3, e2)) & ~ (op(e1, e2) = op(e3, e2)) & ~ (op(e1, e2) = op(e2, e2)) & ~ (op(e0, e2) = op(e3, e2)) & ~ (op(e0, e2) = op(e2, e2)) & ~ (op(e0, e2) = op(e1, e2)) & ~ (op(e2, e1) = op(e3, e1)) & ~ (op(e1, e1) = op(e3, e1)) & ~ (op(e1, e1) = op(e2, e1)) & ~ (op(e0, e1) = op(e3, e1)) & ~ (op(e0, e1) = op(e2, e1)) & ~ (op(e0, e1) = op(e1, e1)) & ~ (op(e2, e0) = op(e3, e0)) & ~ (op(e1, e0) = op(e3, e0)) & ~ (op(e1, e0) = op(e2, e0)) & ~ (op(e0, e0) = op(e3, e0)) & ~ (op(e0, e0) = op(e2, e0)) & ~ (op(e0, e0) = op(e1, e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG154+1.p', ax3)).
fof(f4125, plain, (spl66_5 | ~ spl66_15 | ~ spl66_146), inference(avatar_contradiction_clause, [], [f4124])).
fof(f4124, plain, ($false | (spl66_5 | ~ spl66_15 | ~ spl66_146)), inference(subsumption_resolution, [], [f4122, f620])).
fof(f620, plain, (~ (e0 = op(e3, e2)) | spl66_5), inference(avatar_component_clause, [], [f619])).
fof(f619, plain, (spl66_5 <=> (e0 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_5])])).
fof(f4122, plain, ((e0 = op(e3, e2)) | (~ spl66_15 | ~ spl66_146)), inference(backward_demodulation, [], [f1511, f663])).
fof(f663, plain, ((e2 = op(e3, e0)) | ~ spl66_15), inference(avatar_component_clause, [], [f661])).
fof(f661, plain, (spl66_15 <=> (e2 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_15])])).
fof(f1511, plain, ((e0 = op(e3, op(e3, e0))) | ~ spl66_146), inference(avatar_component_clause, [], [f1509])).
fof(f1509, plain, (spl66_146 <=> (e0 = op(e3, op(e3, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl66_146])])).
fof(f4114, plain, (~ spl66_6 | ~ spl66_24 | spl66_160), inference(avatar_split_clause, [], [f4109, f1582, f699, f623])).
fof(f623, plain, (spl66_6 <=> (e1 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_6])])).
fof(f699, plain, (spl66_24 <=> (e3 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_24])])).
fof(f1582, plain, (spl66_160 <=> (e1 = op(op(e2, e2), e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_160])])).
fof(f4109, plain, (~ (e1 = op(e3, e2)) | (~ spl66_24 | spl66_160)), inference(backward_demodulation, [], [f1584, f701])).
fof(f701, plain, ((e3 = op(e2, e2)) | ~ spl66_24), inference(avatar_component_clause, [], [f699])).
fof(f1584, plain, (~ (e1 = op(op(e2, e2), e2)) | spl66_160), inference(avatar_component_clause, [], [f1582])).
fof(f4112, plain, (spl66_19 | ~ spl66_24 | ~ spl66_67), inference(avatar_split_clause, [], [f4107, f918, f699, f678])).
fof(f678, plain, (spl66_19 <=> (e2 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_19])])).
fof(f918, plain, (spl66_67 <=> (e2 = op(e2, op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl66_67])])).
fof(f4107, plain, ((e2 = op(e2, e3)) | (~ spl66_24 | ~ spl66_67)), inference(backward_demodulation, [], [f920, f701])).
fof(f920, plain, ((e2 = op(e2, op(e2, e2))) | ~ spl66_67), inference(avatar_component_clause, [], [f918])).
fof(f4096, plain, (~ spl66_21 | ~ spl66_29), inference(avatar_split_clause, [], [f4089, f721, f687])).
fof(f687, plain, (spl66_21 <=> (e0 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_21])])).
fof(f721, plain, (spl66_29 <=> (e0 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_29])])).
fof(f4089, plain, (~ (e0 = op(e2, e2)) | ~ spl66_29), inference(backward_demodulation, [], [f272, f723])).
fof(f723, plain, ((e0 = op(e2, e0)) | ~ spl66_29), inference(avatar_component_clause, [], [f721])).
fof(f272, plain, ~ (op(e2, e0) = op(e2, e2)), inference(cnf_transformation, [], [f3])).
fof(f4086, plain, (~ spl66_18 | ~ spl66_34), inference(avatar_split_clause, [], [f4082, f742, f674])).
fof(f674, plain, (spl66_18 <=> (e1 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_18])])).
fof(f742, plain, (spl66_34 <=> (e1 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_34])])).
fof(f4082, plain, (~ (e1 = op(e2, e3)) | ~ spl66_34), inference(backward_demodulation, [], [f256, f744])).
fof(f744, plain, ((e1 = op(e1, e3)) | ~ spl66_34), inference(avatar_component_clause, [], [f742])).
fof(f256, plain, ~ (op(e1, e3) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f4081, plain, (~ spl66_7 | ~ spl66_39), inference(avatar_split_clause, [], [f4078, f763, f627])).
fof(f627, plain, (spl66_7 <=> (e2 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_7])])).
fof(f763, plain, (spl66_39 <=> (e2 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_39])])).
fof(f4078, plain, (~ (e2 = op(e3, e2)) | ~ spl66_39), inference(backward_demodulation, [], [f251, f765])).
fof(f765, plain, ((e2 = op(e1, e2)) | ~ spl66_39), inference(avatar_component_clause, [], [f763])).
fof(f251, plain, ~ (op(e1, e2) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f4075, plain, (~ spl66_40 | ~ spl66_48), inference(avatar_split_clause, [], [f4071, f801, f767])).
fof(f767, plain, (spl66_40 <=> (e3 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_40])])).
fof(f801, plain, (spl66_48 <=> (e3 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_48])])).
fof(f4071, plain, (~ (e3 = op(e1, e2)) | ~ spl66_48), inference(backward_demodulation, [], [f266, f803])).
fof(f803, plain, ((e3 = op(e1, e0)) | ~ spl66_48), inference(avatar_component_clause, [], [f801])).
fof(f266, plain, ~ (op(e1, e0) = op(e1, e2)), inference(cnf_transformation, [], [f3])).
fof(f4074, plain, (~ spl66_32 | ~ spl66_48), inference(avatar_split_clause, [], [f4069, f801, f733])).
fof(f733, plain, (spl66_32 <=> (e3 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_32])])).
fof(f4069, plain, (~ (e3 = op(e2, e0)) | ~ spl66_48), inference(backward_demodulation, [], [f238, f803])).
fof(f238, plain, ~ (op(e1, e0) = op(e2, e0)), inference(cnf_transformation, [], [f3])).
fof(f4068, plain, (~ spl66_4 | ~ spl66_52), inference(avatar_split_clause, [], [f4066, f818, f614])).
fof(f818, plain, (spl66_52 <=> (e3 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_52])])).
fof(f4066, plain, (~ (e3 = op(e3, e3)) | ~ spl66_52), inference(backward_demodulation, [], [f255, f820])).
fof(f820, plain, ((e3 = op(e0, e3)) | ~ spl66_52), inference(avatar_component_clause, [], [f818])).
fof(f255, plain, ~ (op(e0, e3) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f4062, plain, (~ spl66_51 | ~ spl66_59), inference(avatar_split_clause, [], [f4058, f848, f814])).
fof(f814, plain, (spl66_51 <=> (e2 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_51])])).
fof(f848, plain, (spl66_59 <=> (e2 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_59])])).
fof(f4058, plain, (~ (e2 = op(e0, e3)) | ~ spl66_59), inference(backward_demodulation, [], [f263, f850])).
fof(f850, plain, ((e2 = op(e0, e1)) | ~ spl66_59), inference(avatar_component_clause, [], [f848])).
fof(f263, plain, ~ (op(e0, e1) = op(e0, e3)), inference(cnf_transformation, [], [f3])).
fof(f4061, plain, (~ spl66_27 | ~ spl66_59), inference(avatar_split_clause, [], [f4056, f848, f712])).
fof(f712, plain, (spl66_27 <=> (e2 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_27])])).
fof(f4056, plain, (~ (e2 = op(e2, e1)) | ~ spl66_59), inference(backward_demodulation, [], [f242, f850])).
fof(f242, plain, ~ (op(e0, e1) = op(e2, e1)), inference(cnf_transformation, [], [f3])).
fof(f4055, plain, (~ spl66_47 | ~ spl66_62 | spl66_148), inference(avatar_split_clause, [], [f4054, f1522, f861, f797])).
fof(f797, plain, (spl66_47 <=> (e2 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_47])])).
fof(f861, plain, (spl66_62 <=> (op(e0, e0) = e1)), introduced(avatar_definition, [new_symbols(naming, [spl66_62])])).
fof(f1522, plain, (spl66_148 <=> (e2 = op(op(e0, e0), e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_148])])).
fof(f4054, plain, (~ (e2 = op(e1, e0)) | (~ spl66_62 | spl66_148)), inference(forward_demodulation, [], [f1524, f863])).
fof(f863, plain, ((op(e0, e0) = e1) | ~ spl66_62), inference(avatar_component_clause, [], [f861])).
fof(f1524, plain, (~ (e2 = op(op(e0, e0), e0)) | spl66_148), inference(avatar_component_clause, [], [f1522])).
fof(f4049, plain, (~ spl66_60 | ~ spl66_41 | spl66_155), inference(avatar_split_clause, [], [f4048, f1557, f772, f852])).
fof(f852, plain, (spl66_60 <=> (e3 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_60])])).
fof(f772, plain, (spl66_41 <=> (e0 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_41])])).
fof(f1557, plain, (spl66_155 <=> (e3 = op(op(e1, e1), e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_155])])).
fof(f4048, plain, (~ (e3 = op(e0, e1)) | (~ spl66_41 | spl66_155)), inference(forward_demodulation, [], [f1559, f774])).
fof(f774, plain, ((e0 = op(e1, e1)) | ~ spl66_41), inference(avatar_component_clause, [], [f772])).
fof(f1559, plain, (~ (e3 = op(op(e1, e1), e1)) | spl66_155), inference(avatar_component_clause, [], [f1557])).
fof(f4041, plain, (~ spl66_21 | ~ spl66_53), inference(avatar_split_clause, [], [f4040, f823, f687])).
fof(f823, plain, (spl66_53 <=> (e0 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_53])])).
fof(f4040, plain, (~ (e0 = op(e2, e2)) | ~ spl66_53), inference(forward_demodulation, [], [f248, f825])).
fof(f825, plain, ((e0 = op(e0, e2)) | ~ spl66_53), inference(avatar_component_clause, [], [f823])).
fof(f248, plain, ~ (op(e0, e2) = op(e2, e2)), inference(cnf_transformation, [], [f3])).
fof(f4039, plain, (spl66_3 | ~ spl66_4 | ~ spl66_170), inference(avatar_contradiction_clause, [], [f4038])).
fof(f4038, plain, ($false | (spl66_3 | ~ spl66_4 | ~ spl66_170)), inference(subsumption_resolution, [], [f4036, f611])).
fof(f611, plain, (~ (e2 = op(e3, e3)) | spl66_3), inference(avatar_component_clause, [], [f610])).
fof(f610, plain, (spl66_3 <=> (e2 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_3])])).
fof(f4036, plain, ((e2 = op(e3, e3)) | (~ spl66_4 | ~ spl66_170)), inference(backward_demodulation, [], [f1633, f616])).
fof(f616, plain, ((e3 = op(e3, e3)) | ~ spl66_4), inference(avatar_component_clause, [], [f614])).
fof(f1633, plain, ((e2 = op(op(e3, e3), op(e3, e3))) | ~ spl66_170), inference(avatar_component_clause, [], [f1632])).
fof(f1632, plain, (spl66_170 <=> (e2 = op(op(e3, e3), op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl66_170])])).
fof(f4037, plain, (~ spl66_1 | ~ spl66_4 | spl66_169), inference(avatar_split_clause, [], [f4035, f1626, f614, f602])).
fof(f602, plain, (spl66_1 <=> (e0 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_1])])).
fof(f1626, plain, (spl66_169 <=> (e0 = op(op(e3, e3), e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_169])])).
fof(f4035, plain, (~ (e0 = op(e3, e3)) | (~ spl66_4 | spl66_169)), inference(backward_demodulation, [], [f1628, f616])).
fof(f1628, plain, (~ (e0 = op(op(e3, e3), e3)) | spl66_169), inference(avatar_component_clause, [], [f1626])).
fof(f4025, plain, (~ spl66_1 | ~ spl66_13), inference(avatar_split_clause, [], [f4020, f653, f602])).
fof(f653, plain, (spl66_13 <=> (e0 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_13])])).
fof(f4020, plain, (~ (e0 = op(e3, e3)) | ~ spl66_13), inference(backward_demodulation, [], [f279, f655])).
fof(f655, plain, ((e0 = op(e3, e0)) | ~ spl66_13), inference(avatar_component_clause, [], [f653])).
fof(f279, plain, ~ (op(e3, e0) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f4011, plain, (~ spl66_19 | ~ spl66_27), inference(avatar_split_clause, [], [f4007, f712, f678])).
fof(f4007, plain, (~ (e2 = op(e2, e3)) | ~ spl66_27), inference(backward_demodulation, [], [f275, f714])).
fof(f714, plain, ((e2 = op(e2, e1)) | ~ spl66_27), inference(avatar_component_clause, [], [f712])).
fof(f275, plain, ~ (op(e2, e1) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f4010, plain, (~ spl66_11 | ~ spl66_27), inference(avatar_split_clause, [], [f4006, f712, f644])).
fof(f644, plain, (spl66_11 <=> (e2 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_11])])).
fof(f4006, plain, (~ (e2 = op(e3, e1)) | ~ spl66_27), inference(backward_demodulation, [], [f246, f714])).
fof(f246, plain, ~ (op(e2, e1) = op(e3, e1)), inference(cnf_transformation, [], [f3])).
fof(f4005, plain, (~ spl66_1 | ~ spl66_33), inference(avatar_split_clause, [], [f4001, f738, f602])).
fof(f738, plain, (spl66_33 <=> (e0 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_33])])).
fof(f4001, plain, (~ (e0 = op(e3, e3)) | ~ spl66_33), inference(backward_demodulation, [], [f257, f740])).
fof(f740, plain, ((e0 = op(e1, e3)) | ~ spl66_33), inference(avatar_component_clause, [], [f738])).
fof(f257, plain, ~ (op(e1, e3) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f3999, plain, (~ spl66_38 | ~ spl66_40), inference(avatar_contradiction_clause, [], [f3998])).
fof(f3998, plain, ($false | (~ spl66_38 | ~ spl66_40)), inference(subsumption_resolution, [], [f3997, f287])).
fof(f287, plain, ~ (e1 = e3), inference(cnf_transformation, [], [f4])).
fof(f4, plain, (~ (e2 = e3) & ~ (e1 = e3) & ~ (e1 = e2) & ~ (e0 = e3) & ~ (e0 = e2) & ~ (e0 = e1)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG154+1.p', ax4)).
fof(f3997, plain, ((e1 = e3) | (~ spl66_38 | ~ spl66_40)), inference(backward_demodulation, [], [f769, f761])).
fof(f761, plain, ((e1 = op(e1, e2)) | ~ spl66_38), inference(avatar_component_clause, [], [f759])).
fof(f759, plain, (spl66_38 <=> (e1 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_38])])).
fof(f769, plain, ((e3 = op(e1, e2)) | ~ spl66_40), inference(avatar_component_clause, [], [f767])).
fof(f3995, plain, (~ spl66_8 | ~ spl66_40), inference(avatar_split_clause, [], [f3992, f767, f631])).
fof(f631, plain, (spl66_8 <=> (e3 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_8])])).
fof(f3992, plain, (~ (e3 = op(e3, e2)) | ~ spl66_40), inference(backward_demodulation, [], [f251, f769])).
fof(f3989, plain, (~ spl66_33 | ~ spl66_41), inference(avatar_split_clause, [], [f3980, f772, f738])).
fof(f3980, plain, (~ (e0 = op(e1, e3)) | ~ spl66_41), inference(backward_demodulation, [], [f269, f774])).
fof(f269, plain, ~ (op(e1, e1) = op(e1, e3)), inference(cnf_transformation, [], [f3])).
fof(f3976, plain, (~ spl66_39 | ~ spl66_47), inference(avatar_split_clause, [], [f3972, f797, f763])).
fof(f3972, plain, (~ (e2 = op(e1, e2)) | ~ spl66_47), inference(backward_demodulation, [], [f266, f799])).
fof(f799, plain, ((e2 = op(e1, e0)) | ~ spl66_47), inference(avatar_component_clause, [], [f797])).
fof(f3968, plain, (~ spl66_19 | ~ spl66_51), inference(avatar_split_clause, [], [f3964, f814, f678])).
fof(f3964, plain, (~ (e2 = op(e2, e3)) | ~ spl66_51), inference(backward_demodulation, [], [f254, f816])).
fof(f816, plain, ((e2 = op(e0, e3)) | ~ spl66_51), inference(avatar_component_clause, [], [f814])).
fof(f254, plain, ~ (op(e0, e3) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f3956, plain, (~ spl66_52 | ~ spl66_60), inference(avatar_split_clause, [], [f3954, f852, f818])).
fof(f3954, plain, (~ (e3 = op(e0, e3)) | ~ spl66_60), inference(backward_demodulation, [], [f263, f854])).
fof(f854, plain, ((e3 = op(e0, e1)) | ~ spl66_60), inference(avatar_component_clause, [], [f852])).
fof(f3955, plain, (~ spl66_44 | ~ spl66_60), inference(avatar_split_clause, [], [f3950, f852, f784])).
fof(f784, plain, (spl66_44 <=> (e3 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_44])])).
fof(f3950, plain, (~ (e3 = op(e1, e1)) | ~ spl66_60), inference(backward_demodulation, [], [f241, f854])).
fof(f241, plain, ~ (op(e0, e1) = op(e1, e1)), inference(cnf_transformation, [], [f3])).
fof(f3949, plain, (~ spl66_44 | ~ spl66_62 | spl66_152), inference(avatar_split_clause, [], [f3944, f1542, f861, f784])).
fof(f1542, plain, (spl66_152 <=> (e3 = op(op(e0, e0), op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl66_152])])).
fof(f3944, plain, (~ (e3 = op(e1, e1)) | (~ spl66_62 | spl66_152)), inference(backward_demodulation, [], [f1544, f863])).
fof(f1544, plain, (~ (e3 = op(op(e0, e0), op(e0, e0))) | spl66_152), inference(avatar_component_clause, [], [f1542])).
fof(f3945, plain, (~ spl66_14 | ~ spl66_62), inference(avatar_split_clause, [], [f3935, f861, f657])).
fof(f657, plain, (spl66_14 <=> (e1 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_14])])).
fof(f3935, plain, (~ (e1 = op(e3, e0)) | ~ spl66_62), inference(backward_demodulation, [], [f237, f863])).
fof(f237, plain, ~ (op(e0, e0) = op(e3, e0)), inference(cnf_transformation, [], [f3])).
fof(f3932, plain, (~ spl66_54 | ~ spl66_21 | spl66_160), inference(avatar_split_clause, [], [f3931, f1582, f687, f827])).
fof(f827, plain, (spl66_54 <=> (e1 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_54])])).
fof(f3931, plain, (~ (e1 = op(e0, e2)) | (~ spl66_21 | spl66_160)), inference(forward_demodulation, [], [f1584, f689])).
fof(f689, plain, ((e0 = op(e2, e2)) | ~ spl66_21), inference(avatar_component_clause, [], [f687])).
fof(f3925, plain, (~ spl66_56 | ~ spl66_21 | spl66_161), inference(avatar_split_clause, [], [f3856, f1587, f687, f835])).
fof(f835, plain, (spl66_56 <=> (e3 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_56])])).
fof(f1587, plain, (spl66_161 <=> (e3 = op(op(e2, e2), e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_161])])).
fof(f3856, plain, (~ (e3 = op(e0, e2)) | (~ spl66_21 | spl66_161)), inference(forward_demodulation, [], [f1589, f689])).
fof(f1589, plain, (~ (e3 = op(op(e2, e2), e2)) | spl66_161), inference(avatar_component_clause, [], [f1587])).
fof(f3920, plain, (~ spl66_1 | ~ spl66_14 | ~ spl66_143), inference(avatar_contradiction_clause, [], [f3919])).
fof(f3919, plain, ($false | (~ spl66_1 | ~ spl66_14 | ~ spl66_143)), inference(subsumption_resolution, [], [f3918, f287])).
fof(f3918, plain, ((e1 = e3) | (~ spl66_1 | ~ spl66_14 | ~ spl66_143)), inference(forward_demodulation, [], [f3912, f659])).
fof(f659, plain, ((e1 = op(e3, e0)) | ~ spl66_14), inference(avatar_component_clause, [], [f657])).
fof(f3912, plain, ((e3 = op(e3, e0)) | (~ spl66_1 | ~ spl66_143)), inference(backward_demodulation, [], [f1496, f604])).
fof(f604, plain, ((e0 = op(e3, e3)) | ~ spl66_1), inference(avatar_component_clause, [], [f602])).
fof(f1496, plain, ((e3 = op(e3, op(e3, e3))) | ~ spl66_143), inference(avatar_component_clause, [], [f1494])).
fof(f1494, plain, (spl66_143 <=> (e3 = op(e3, op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl66_143])])).
fof(f3906, plain, (~ spl66_7 | ~ spl66_11), inference(avatar_split_clause, [], [f3902, f644, f627])).
fof(f3902, plain, (~ (e2 = op(e3, e2)) | ~ spl66_11), inference(backward_demodulation, [], [f280, f646])).
fof(f646, plain, ((e2 = op(e3, e1)) | ~ spl66_11), inference(avatar_component_clause, [], [f644])).
fof(f280, plain, ~ (op(e3, e1) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f3894, plain, (~ spl66_2 | ~ spl66_34), inference(avatar_split_clause, [], [f3890, f742, f606])).
fof(f606, plain, (spl66_2 <=> (e1 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_2])])).
fof(f3890, plain, (~ (e1 = op(e3, e3)) | ~ spl66_34), inference(backward_demodulation, [], [f257, f744])).
fof(f3883, plain, (~ spl66_1 | ~ spl66_44 | spl66_153), inference(avatar_split_clause, [], [f3875, f1548, f784, f602])).
fof(f1548, plain, (spl66_153 <=> (e0 = op(op(e1, e1), op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl66_153])])).
fof(f3875, plain, (~ (e0 = op(e3, e3)) | (~ spl66_44 | spl66_153)), inference(backward_demodulation, [], [f1550, f786])).
fof(f786, plain, ((e3 = op(e1, e1)) | ~ spl66_44), inference(avatar_component_clause, [], [f784])).
fof(f1550, plain, (~ (e0 = op(op(e1, e1), op(e1, e1))) | spl66_153), inference(avatar_component_clause, [], [f1548])).
fof(f3882, plain, (spl66_34 | ~ spl66_44 | ~ spl66_73), inference(avatar_split_clause, [], [f3874, f950, f784, f742])).
fof(f950, plain, (spl66_73 <=> (e1 = op(e1, op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl66_73])])).
fof(f3874, plain, ((e1 = op(e1, e3)) | (~ spl66_44 | ~ spl66_73)), inference(backward_demodulation, [], [f952, f786])).
fof(f952, plain, ((e1 = op(e1, op(e1, e1))) | ~ spl66_73), inference(avatar_component_clause, [], [f950])).
fof(f3881, plain, (~ spl66_40 | ~ spl66_44), inference(avatar_split_clause, [], [f3872, f784, f767])).
fof(f3872, plain, (~ (e3 = op(e1, e2)) | ~ spl66_44), inference(backward_demodulation, [], [f268, f786])).
fof(f268, plain, ~ (op(e1, e1) = op(e1, e2)), inference(cnf_transformation, [], [f3])).
fof(f3880, plain, (~ spl66_12 | ~ spl66_44), inference(avatar_split_clause, [], [f3871, f784, f648])).
fof(f3871, plain, (~ (e3 = op(e3, e1)) | ~ spl66_44), inference(backward_demodulation, [], [f245, f786])).
fof(f245, plain, ~ (op(e1, e1) = op(e3, e1)), inference(cnf_transformation, [], [f3])).
fof(f3870, plain, (~ spl66_33 | ~ spl66_45), inference(avatar_split_clause, [], [f3864, f789, f738])).
fof(f789, plain, (spl66_45 <=> (e0 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_45])])).
fof(f3864, plain, (~ (e0 = op(e1, e3)) | ~ spl66_45), inference(backward_demodulation, [], [f267, f791])).
fof(f791, plain, ((e0 = op(e1, e0)) | ~ spl66_45), inference(avatar_component_clause, [], [f789])).
fof(f267, plain, ~ (op(e1, e0) = op(e1, e3)), inference(cnf_transformation, [], [f3])).
fof(f3869, plain, (~ spl66_13 | ~ spl66_45), inference(avatar_split_clause, [], [f3861, f789, f653])).
fof(f3861, plain, (~ (e0 = op(e3, e0)) | ~ spl66_45), inference(backward_demodulation, [], [f239, f791])).
fof(f239, plain, ~ (op(e1, e0) = op(e3, e0)), inference(cnf_transformation, [], [f3])).
fof(f3849, plain, (~ spl66_2 | ~ spl66_43 | spl66_170), inference(avatar_contradiction_clause, [], [f3848])).
fof(f3848, plain, ($false | (~ spl66_2 | ~ spl66_43 | spl66_170)), inference(subsumption_resolution, [], [f3847, f782])).
fof(f782, plain, ((e2 = op(e1, e1)) | ~ spl66_43), inference(avatar_component_clause, [], [f780])).
fof(f780, plain, (spl66_43 <=> (e2 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_43])])).
fof(f3847, plain, (~ (e2 = op(e1, e1)) | (~ spl66_2 | spl66_170)), inference(backward_demodulation, [], [f1634, f608])).
fof(f608, plain, ((e1 = op(e3, e3)) | ~ spl66_2), inference(avatar_component_clause, [], [f606])).
fof(f1634, plain, (~ (e2 = op(op(e3, e3), op(e3, e3))) | spl66_170), inference(avatar_component_clause, [], [f1632])).
fof(f3841, plain, (spl66_2 | ~ spl66_12 | ~ spl66_145), inference(avatar_split_clause, [], [f3839, f1504, f648, f606])).
fof(f1504, plain, (spl66_145 <=> (e1 = op(e3, op(e3, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl66_145])])).
fof(f3839, plain, ((e1 = op(e3, e3)) | (~ spl66_12 | ~ spl66_145)), inference(backward_demodulation, [], [f1506, f650])).
fof(f1506, plain, ((e1 = op(e3, op(e3, e1))) | ~ spl66_145), inference(avatar_component_clause, [], [f1504])).
fof(f3813, plain, (~ spl66_28 | ~ spl66_32), inference(avatar_split_clause, [], [f3810, f733, f716])).
fof(f716, plain, (spl66_28 <=> (e3 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_28])])).
fof(f3810, plain, (~ (e3 = op(e2, e1)) | ~ spl66_32), inference(backward_demodulation, [], [f271, f735])).
fof(f735, plain, ((e3 = op(e2, e0)) | ~ spl66_32), inference(avatar_component_clause, [], [f733])).
fof(f271, plain, ~ (op(e2, e0) = op(e2, e1)), inference(cnf_transformation, [], [f3])).
fof(f3802, plain, (~ spl66_35 | ~ spl66_40 | spl66_72), inference(avatar_split_clause, [], [f3801, f945, f767, f746])).
fof(f746, plain, (spl66_35 <=> (e2 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_35])])).
fof(f945, plain, (spl66_72 <=> (e2 = op(e1, op(e1, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl66_72])])).
fof(f3801, plain, (~ (e2 = op(e1, e3)) | (~ spl66_40 | spl66_72)), inference(backward_demodulation, [], [f946, f769])).
fof(f946, plain, (~ (e2 = op(e1, op(e1, e2))) | spl66_72), inference(avatar_component_clause, [], [f945])).
fof(f3798, plain, (~ spl66_14 | ~ spl66_46), inference(avatar_split_clause, [], [f3792, f793, f657])).
fof(f793, plain, (spl66_46 <=> (e1 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_46])])).
fof(f3792, plain, (~ (e1 = op(e3, e0)) | ~ spl66_46), inference(backward_demodulation, [], [f239, f795])).
fof(f795, plain, ((e1 = op(e1, e0)) | ~ spl66_46), inference(avatar_component_clause, [], [f793])).
fof(f3790, plain, (~ spl66_52 | spl66_76), inference(avatar_contradiction_clause, [], [f3789])).
fof(f3789, plain, ($false | (~ spl66_52 | spl66_76)), inference(subsumption_resolution, [], [f3787, f820])).
fof(f3787, plain, (~ (e3 = op(e0, e3)) | (~ spl66_52 | spl66_76)), inference(backward_demodulation, [], [f968, f820])).
fof(f968, plain, (~ (e3 = op(e0, op(e0, e3))) | spl66_76), inference(avatar_component_clause, [], [f967])).
fof(f967, plain, (spl66_76 <=> (e3 = op(e0, op(e0, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl66_76])])).
fof(f3783, plain, (~ spl66_50 | ~ spl66_54), inference(avatar_split_clause, [], [f3778, f827, f810])).
fof(f810, plain, (spl66_50 <=> (e1 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_50])])).
fof(f3778, plain, (~ (e1 = op(e0, e3)) | ~ spl66_54), inference(backward_demodulation, [], [f264, f829])).
fof(f829, plain, ((e1 = op(e0, e2)) | ~ spl66_54), inference(avatar_component_clause, [], [f827])).
fof(f264, plain, ~ (op(e0, e2) = op(e0, e3)), inference(cnf_transformation, [], [f3])).
fof(f3782, plain, (~ spl66_22 | ~ spl66_54), inference(avatar_split_clause, [], [f3777, f827, f691])).
fof(f691, plain, (spl66_22 <=> (e1 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_22])])).
fof(f3777, plain, (~ (e1 = op(e2, e2)) | ~ spl66_54), inference(backward_demodulation, [], [f248, f829])).
fof(f3770, plain, (~ spl66_160 | ~ spl66_163), inference(avatar_contradiction_clause, [], [f3769])).
fof(f3769, plain, ($false | (~ spl66_160 | ~ spl66_163)), inference(subsumption_resolution, [], [f3768, f283])).
fof(f283, plain, ~ (e0 = e1), inference(cnf_transformation, [], [f4])).
fof(f3768, plain, ((e0 = e1) | (~ spl66_160 | ~ spl66_163)), inference(forward_demodulation, [], [f1583, f1597])).
fof(f1597, plain, ((e0 = op(op(e2, e2), e2)) | ~ spl66_163), inference(avatar_component_clause, [], [f1596])).
fof(f1596, plain, (spl66_163 <=> (e0 = op(op(e2, e2), e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_163])])).
fof(f1583, plain, ((e1 = op(op(e2, e2), e2)) | ~ spl66_160), inference(avatar_component_clause, [], [f1582])).
fof(f3754, plain, (~ spl66_24 | ~ spl66_43 | spl66_158), inference(avatar_split_clause, [], [f3686, f1572, f780, f699])).
fof(f1572, plain, (spl66_158 <=> (e3 = op(op(e1, e1), op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl66_158])])).
fof(f3686, plain, (~ (e3 = op(e2, e2)) | (~ spl66_43 | spl66_158)), inference(backward_demodulation, [], [f1574, f782])).
fof(f1574, plain, (~ (e3 = op(op(e1, e1), op(e1, e1))) | spl66_158), inference(avatar_component_clause, [], [f1572])).
fof(f3753, plain, (~ spl66_21 | ~ spl66_43 | spl66_153), inference(avatar_split_clause, [], [f3682, f1548, f780, f687])).
fof(f3682, plain, (~ (e0 = op(e2, e2)) | (~ spl66_43 | spl66_153)), inference(backward_demodulation, [], [f1550, f782])).
fof(f3748, plain, (~ spl66_31 | ~ spl66_19), inference(avatar_split_clause, [], [f3747, f678, f729])).
fof(f729, plain, (spl66_31 <=> (e2 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_31])])).
fof(f3747, plain, (~ (e2 = op(e2, e0)) | ~ spl66_19), inference(forward_demodulation, [], [f273, f680])).
fof(f680, plain, ((e2 = op(e2, e3)) | ~ spl66_19), inference(avatar_component_clause, [], [f678])).
fof(f273, plain, ~ (op(e2, e0) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f3745, plain, (~ spl66_35 | ~ spl66_19), inference(avatar_split_clause, [], [f3744, f678, f746])).
fof(f3744, plain, (~ (e2 = op(e1, e3)) | ~ spl66_19), inference(forward_demodulation, [], [f256, f680])).
fof(f3743, plain, (~ spl66_23 | ~ spl66_19), inference(avatar_split_clause, [], [f3742, f678, f695])).
fof(f695, plain, (spl66_23 <=> (e2 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_23])])).
fof(f3742, plain, (~ (e2 = op(e2, e2)) | ~ spl66_19), inference(forward_demodulation, [], [f276, f680])).
fof(f276, plain, ~ (op(e2, e2) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f3736, plain, (spl66_1 | ~ spl66_4 | ~ spl66_169), inference(avatar_contradiction_clause, [], [f3735])).
fof(f3735, plain, ($false | (spl66_1 | ~ spl66_4 | ~ spl66_169)), inference(subsumption_resolution, [], [f3734, f603])).
fof(f603, plain, (~ (e0 = op(e3, e3)) | spl66_1), inference(avatar_component_clause, [], [f602])).
fof(f3734, plain, ((e0 = op(e3, e3)) | (~ spl66_4 | ~ spl66_169)), inference(backward_demodulation, [], [f1627, f616])).
fof(f1627, plain, ((e0 = op(op(e3, e3), e3)) | ~ spl66_169), inference(avatar_component_clause, [], [f1626])).
fof(f3731, plain, (~ spl66_3 | ~ spl66_7), inference(avatar_split_clause, [], [f3728, f627, f610])).
fof(f3728, plain, (~ (e2 = op(e3, e3)) | ~ spl66_7), inference(backward_demodulation, [], [f282, f629])).
fof(f629, plain, ((e2 = op(e3, e2)) | ~ spl66_7), inference(avatar_component_clause, [], [f627])).
fof(f282, plain, ~ (op(e3, e2) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f3721, plain, (spl66_9 | ~ spl66_14 | ~ spl66_146), inference(avatar_split_clause, [], [f3719, f1509, f657, f636])).
fof(f636, plain, (spl66_9 <=> (e0 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_9])])).
fof(f3719, plain, ((e0 = op(e3, e1)) | (~ spl66_14 | ~ spl66_146)), inference(backward_demodulation, [], [f1511, f659])).
fof(f3720, plain, (~ spl66_10 | ~ spl66_14), inference(avatar_split_clause, [], [f3715, f657, f640])).
fof(f640, plain, (spl66_10 <=> (e1 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_10])])).
fof(f3715, plain, (~ (e1 = op(e3, e1)) | ~ spl66_14), inference(backward_demodulation, [], [f277, f659])).
fof(f277, plain, ~ (op(e3, e0) = op(e3, e1)), inference(cnf_transformation, [], [f3])).
fof(f3714, plain, (~ spl66_17 | ~ spl66_19), inference(avatar_contradiction_clause, [], [f3713])).
fof(f3713, plain, ($false | (~ spl66_17 | ~ spl66_19)), inference(subsumption_resolution, [], [f3712, f284])).
fof(f284, plain, ~ (e0 = e2), inference(cnf_transformation, [], [f4])).
fof(f3712, plain, ((e0 = e2) | (~ spl66_17 | ~ spl66_19)), inference(forward_demodulation, [], [f680, f672])).
fof(f672, plain, ((e0 = op(e2, e3)) | ~ spl66_17), inference(avatar_component_clause, [], [f670])).
fof(f670, plain, (spl66_17 <=> (e0 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_17])])).
fof(f3704, plain, (~ spl66_3 | ~ spl66_35), inference(avatar_split_clause, [], [f3701, f746, f610])).
fof(f3701, plain, (~ (e2 = op(e3, e3)) | ~ spl66_35), inference(backward_demodulation, [], [f257, f748])).
fof(f748, plain, ((e2 = op(e1, e3)) | ~ spl66_35), inference(avatar_component_clause, [], [f746])).
fof(f3688, plain, (~ spl66_35 | ~ spl66_43), inference(avatar_split_clause, [], [f3679, f780, f746])).
fof(f3679, plain, (~ (e2 = op(e1, e3)) | ~ spl66_43), inference(backward_demodulation, [], [f269, f782])).
fof(f3687, plain, (~ spl66_39 | ~ spl66_43), inference(avatar_split_clause, [], [f3678, f780, f763])).
fof(f3678, plain, (~ (e2 = op(e1, e2)) | ~ spl66_43), inference(backward_demodulation, [], [f268, f782])).
fof(f3675, plain, (~ spl66_34 | ~ spl66_50), inference(avatar_split_clause, [], [f3670, f810, f742])).
fof(f3670, plain, (~ (e1 = op(e1, e3)) | ~ spl66_50), inference(backward_demodulation, [], [f253, f812])).
fof(f812, plain, ((e1 = op(e0, e3)) | ~ spl66_50), inference(avatar_component_clause, [], [f810])).
fof(f253, plain, ~ (op(e0, e3) = op(e1, e3)), inference(cnf_transformation, [], [f3])).
fof(f3668, plain, (~ spl66_8 | ~ spl66_56), inference(avatar_split_clause, [], [f3665, f835, f631])).
fof(f3665, plain, (~ (e3 = op(e3, e2)) | ~ spl66_56), inference(backward_demodulation, [], [f249, f837])).
fof(f837, plain, ((e3 = op(e0, e2)) | ~ spl66_56), inference(avatar_component_clause, [], [f835])).
fof(f249, plain, ~ (op(e0, e2) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f3663, plain, (~ spl66_53 | ~ spl66_57), inference(avatar_split_clause, [], [f3655, f840, f823])).
fof(f840, plain, (spl66_57 <=> (e0 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_57])])).
fof(f3655, plain, (~ (e0 = op(e0, e2)) | ~ spl66_57), inference(backward_demodulation, [], [f262, f842])).
fof(f842, plain, ((e0 = op(e0, e1)) | ~ spl66_57), inference(avatar_component_clause, [], [f840])).
fof(f262, plain, ~ (op(e0, e1) = op(e0, e2)), inference(cnf_transformation, [], [f3])).
fof(f3662, plain, (~ spl66_9 | ~ spl66_57), inference(avatar_split_clause, [], [f3654, f840, f636])).
fof(f3654, plain, (~ (e0 = op(e3, e1)) | ~ spl66_57), inference(backward_demodulation, [], [f243, f842])).
fof(f243, plain, ~ (op(e0, e1) = op(e3, e1)), inference(cnf_transformation, [], [f3])).
fof(f3661, plain, (~ spl66_41 | ~ spl66_57), inference(avatar_split_clause, [], [f3653, f840, f772])).
fof(f3653, plain, (~ (e0 = op(e1, e1)) | ~ spl66_57), inference(backward_demodulation, [], [f241, f842])).
fof(f3651, plain, (~ spl66_59 | ~ spl66_63), inference(avatar_split_clause, [], [f3641, f865, f848])).
fof(f865, plain, (spl66_63 <=> (op(e0, e0) = e2)), introduced(avatar_definition, [new_symbols(naming, [spl66_63])])).
fof(f3641, plain, (~ (e2 = op(e0, e1)) | ~ spl66_63), inference(backward_demodulation, [], [f259, f867])).
fof(f867, plain, ((op(e0, e0) = e2) | ~ spl66_63), inference(avatar_component_clause, [], [f865])).
fof(f259, plain, ~ (op(e0, e0) = op(e0, e1)), inference(cnf_transformation, [], [f3])).
fof(f3650, plain, (~ spl66_15 | ~ spl66_63), inference(avatar_split_clause, [], [f3640, f865, f661])).
fof(f3640, plain, (~ (e2 = op(e3, e0)) | ~ spl66_63), inference(backward_demodulation, [], [f237, f867])).
fof(f3649, plain, (~ spl66_31 | ~ spl66_63), inference(avatar_split_clause, [], [f3639, f865, f729])).
fof(f3639, plain, (~ (e2 = op(e2, e0)) | ~ spl66_63), inference(backward_demodulation, [], [f236, f867])).
fof(f236, plain, ~ (op(e0, e0) = op(e2, e0)), inference(cnf_transformation, [], [f3])).
fof(f3633, plain, (~ spl66_41 | ~ spl66_22 | spl66_159), inference(avatar_split_clause, [], [f3632, f1578, f691, f772])).
fof(f1578, plain, (spl66_159 <=> (e0 = op(op(e2, e2), op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl66_159])])).
fof(f3632, plain, (~ (e0 = op(e1, e1)) | (~ spl66_22 | spl66_159)), inference(forward_demodulation, [], [f1580, f693])).
fof(f693, plain, ((e1 = op(e2, e2)) | ~ spl66_22), inference(avatar_component_clause, [], [f691])).
fof(f1580, plain, (~ (e0 = op(op(e2, e2), op(e2, e2))) | spl66_159), inference(avatar_component_clause, [], [f1578])).
fof(f3620, plain, (~ spl66_40 | ~ spl66_22 | spl66_161), inference(avatar_split_clause, [], [f3588, f1587, f691, f767])).
fof(f3588, plain, (~ (e3 = op(e1, e2)) | (~ spl66_22 | spl66_161)), inference(backward_demodulation, [], [f1589, f693])).
fof(f3617, plain, (~ spl66_29 | ~ spl66_17), inference(avatar_split_clause, [], [f3334, f670, f721])).
fof(f3334, plain, (~ (e0 = op(e2, e0)) | ~ spl66_17), inference(forward_demodulation, [], [f273, f672])).
fof(f3615, plain, (~ spl66_3 | ~ spl66_22 | spl66_168), inference(avatar_contradiction_clause, [], [f3614])).
fof(f3614, plain, ($false | (~ spl66_3 | ~ spl66_22 | spl66_168)), inference(subsumption_resolution, [], [f3611, f693])).
fof(f3611, plain, (~ (e1 = op(e2, e2)) | (~ spl66_3 | spl66_168)), inference(backward_demodulation, [], [f1624, f612])).
fof(f612, plain, ((e2 = op(e3, e3)) | ~ spl66_3), inference(avatar_component_clause, [], [f610])).
fof(f1624, plain, (~ (e1 = op(op(e3, e3), op(e3, e3))) | spl66_168), inference(avatar_component_clause, [], [f1622])).
fof(f1622, plain, (spl66_168 <=> (e1 = op(op(e3, e3), op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl66_168])])).
fof(f3607, plain, (spl66_3 | ~ spl66_8 | ~ spl66_144), inference(avatar_split_clause, [], [f3606, f1499, f631, f610])).
fof(f1499, plain, (spl66_144 <=> (e2 = op(e3, op(e3, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl66_144])])).
fof(f3606, plain, ((e2 = op(e3, e3)) | (~ spl66_8 | ~ spl66_144)), inference(backward_demodulation, [], [f1501, f633])).
fof(f633, plain, ((e3 = op(e3, e2)) | ~ spl66_8), inference(avatar_component_clause, [], [f631])).
fof(f1501, plain, ((e2 = op(e3, op(e3, e2))) | ~ spl66_144), inference(avatar_component_clause, [], [f1499])).
fof(f3604, plain, (~ spl66_2 | ~ spl66_10), inference(avatar_split_clause, [], [f3600, f640, f606])).
fof(f3600, plain, (~ (e1 = op(e3, e3)) | ~ spl66_10), inference(backward_demodulation, [], [f281, f642])).
fof(f642, plain, ((e1 = op(e3, e1)) | ~ spl66_10), inference(avatar_component_clause, [], [f640])).
fof(f3581, plain, (~ spl66_24 | ~ spl66_28), inference(avatar_split_clause, [], [f3579, f716, f699])).
fof(f3579, plain, (~ (e3 = op(e2, e2)) | ~ spl66_28), inference(backward_demodulation, [], [f274, f718])).
fof(f718, plain, ((e3 = op(e2, e1)) | ~ spl66_28), inference(avatar_component_clause, [], [f716])).
fof(f274, plain, ~ (op(e2, e1) = op(e2, e2)), inference(cnf_transformation, [], [f3])).
fof(f3577, plain, (~ spl66_27 | ~ spl66_31), inference(avatar_split_clause, [], [f3573, f729, f712])).
fof(f3573, plain, (~ (e2 = op(e2, e1)) | ~ spl66_31), inference(backward_demodulation, [], [f271, f731])).
fof(f731, plain, ((e2 = op(e2, e0)) | ~ spl66_31), inference(avatar_component_clause, [], [f729])).
fof(f3556, plain, (~ spl66_9 | ~ spl66_41), inference(avatar_split_clause, [], [f3546, f772, f636])).
fof(f3546, plain, (~ (e0 = op(e3, e1)) | ~ spl66_41), inference(backward_demodulation, [], [f245, f774])).
fof(f3544, plain, (~ spl66_44 | ~ spl66_48), inference(avatar_split_clause, [], [f3539, f801, f784])).
fof(f3539, plain, (~ (e3 = op(e1, e1)) | ~ spl66_48), inference(backward_demodulation, [], [f265, f803])).
fof(f265, plain, ~ (op(e1, e0) = op(e1, e1)), inference(cnf_transformation, [], [f3])).
fof(f3543, plain, (~ spl66_16 | ~ spl66_48), inference(avatar_split_clause, [], [f3538, f801, f665])).
fof(f665, plain, (spl66_16 <=> (e3 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_16])])).
fof(f3538, plain, (~ (e3 = op(e3, e0)) | ~ spl66_48), inference(backward_demodulation, [], [f239, f803])).
fof(f3527, plain, (~ spl66_58 | ~ spl66_62), inference(avatar_split_clause, [], [f3519, f861, f844])).
fof(f844, plain, (spl66_58 <=> (e1 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_58])])).
fof(f3519, plain, (~ (e1 = op(e0, e1)) | ~ spl66_62), inference(backward_demodulation, [], [f259, f863])).
fof(f3526, plain, (~ spl66_30 | ~ spl66_62), inference(avatar_split_clause, [], [f3517, f861, f725])).
fof(f725, plain, (spl66_30 <=> (e1 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_30])])).
fof(f3517, plain, (~ (e1 = op(e2, e0)) | ~ spl66_62), inference(backward_demodulation, [], [f236, f863])).
fof(f3512, plain, (~ spl66_167 | ~ spl66_169), inference(avatar_contradiction_clause, [], [f3511])).
fof(f3511, plain, ($false | (~ spl66_167 | ~ spl66_169)), inference(subsumption_resolution, [], [f3510, f284])).
fof(f3510, plain, ((e0 = e2) | (~ spl66_167 | ~ spl66_169)), inference(backward_demodulation, [], [f1618, f1627])).
fof(f1618, plain, ((e2 = op(op(e3, e3), e3)) | ~ spl66_167), inference(avatar_component_clause, [], [f1617])).
fof(f1617, plain, (spl66_167 <=> (e2 = op(op(e3, e3), e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_167])])).
fof(f3494, plain, (~ spl66_7 | spl66_144), inference(avatar_contradiction_clause, [], [f3493])).
fof(f3493, plain, ($false | (~ spl66_7 | spl66_144)), inference(subsumption_resolution, [], [f3491, f629])).
fof(f3491, plain, (~ (e2 = op(e3, e2)) | (~ spl66_7 | spl66_144)), inference(backward_demodulation, [], [f1500, f629])).
fof(f1500, plain, (~ (e2 = op(e3, op(e3, e2))) | spl66_144), inference(avatar_component_clause, [], [f1499])).
fof(f3482, plain, (~ spl66_12 | ~ spl66_16), inference(avatar_split_clause, [], [f3478, f665, f648])).
fof(f3478, plain, (~ (e3 = op(e3, e1)) | ~ spl66_16), inference(backward_demodulation, [], [f277, f667])).
fof(f667, plain, ((e3 = op(e3, e0)) | ~ spl66_16), inference(avatar_component_clause, [], [f665])).
fof(f3474, plain, (~ spl66_26 | ~ spl66_30), inference(avatar_split_clause, [], [f3470, f725, f708])).
fof(f708, plain, (spl66_26 <=> (e1 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_26])])).
fof(f3470, plain, (~ (e1 = op(e2, e1)) | ~ spl66_30), inference(backward_demodulation, [], [f271, f727])).
fof(f727, plain, ((e1 = op(e2, e0)) | ~ spl66_30), inference(avatar_component_clause, [], [f725])).
fof(f3463, plain, (~ spl66_34 | ~ spl66_38), inference(avatar_split_clause, [], [f3458, f759, f742])).
fof(f3458, plain, (~ (e1 = op(e1, e3)) | ~ spl66_38), inference(backward_demodulation, [], [f270, f761])).
fof(f270, plain, ~ (op(e1, e2) = op(e1, e3)), inference(cnf_transformation, [], [f3])).
fof(f3462, plain, (~ spl66_6 | ~ spl66_38), inference(avatar_split_clause, [], [f3457, f759, f623])).
fof(f3457, plain, (~ (e1 = op(e3, e2)) | ~ spl66_38), inference(backward_demodulation, [], [f251, f761])).
fof(f3455, plain, (~ spl66_3 | ~ spl66_44 | spl66_156), inference(avatar_split_clause, [], [f3452, f1562, f784, f610])).
fof(f1562, plain, (spl66_156 <=> (e2 = op(op(e1, e1), op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl66_156])])).
fof(f3452, plain, (~ (e2 = op(e3, e3)) | (~ spl66_44 | spl66_156)), inference(backward_demodulation, [], [f1564, f786])).
fof(f1564, plain, (~ (e2 = op(op(e1, e1), op(e1, e1))) | spl66_156), inference(avatar_component_clause, [], [f1562])).
fof(f3435, plain, (~ spl66_37 | ~ spl66_53), inference(avatar_split_clause, [], [f3429, f823, f755])).
fof(f755, plain, (spl66_37 <=> (e0 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_37])])).
fof(f3429, plain, (~ (e0 = op(e1, e2)) | ~ spl66_53), inference(backward_demodulation, [], [f247, f825])).
fof(f247, plain, ~ (op(e0, e2) = op(e1, e2)), inference(cnf_transformation, [], [f3])).
fof(f3428, plain, (~ spl66_26 | ~ spl66_58), inference(avatar_split_clause, [], [f3422, f844, f708])).
fof(f3422, plain, (~ (e1 = op(e2, e1)) | ~ spl66_58), inference(backward_demodulation, [], [f242, f846])).
fof(f846, plain, ((e1 = op(e0, e1)) | ~ spl66_58), inference(avatar_component_clause, [], [f844])).
fof(f3420, plain, (~ spl66_24 | ~ spl66_63 | spl66_152), inference(avatar_contradiction_clause, [], [f3419])).
fof(f3419, plain, ($false | (~ spl66_24 | ~ spl66_63 | spl66_152)), inference(subsumption_resolution, [], [f3414, f701])).
fof(f3414, plain, (~ (e3 = op(e2, e2)) | (~ spl66_63 | spl66_152)), inference(backward_demodulation, [], [f1544, f867])).
fof(f3417, plain, (spl66_53 | ~ spl66_63 | ~ spl66_79), inference(avatar_split_clause, [], [f3412, f982, f865, f823])).
fof(f982, plain, (spl66_79 <=> (e0 = op(e0, op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl66_79])])).
fof(f3412, plain, ((e0 = op(e0, e2)) | (~ spl66_63 | ~ spl66_79)), inference(backward_demodulation, [], [f984, f867])).
fof(f984, plain, ((e0 = op(e0, op(e0, e0))) | ~ spl66_79), inference(avatar_component_clause, [], [f982])).
fof(f3416, plain, (~ spl66_55 | ~ spl66_63), inference(avatar_split_clause, [], [f3410, f865, f831])).
fof(f831, plain, (spl66_55 <=> (e2 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_55])])).
fof(f3410, plain, (~ (e2 = op(e0, e2)) | ~ spl66_63), inference(backward_demodulation, [], [f260, f867])).
fof(f260, plain, ~ (op(e0, e0) = op(e0, e2)), inference(cnf_transformation, [], [f3])).
fof(f3395, plain, (~ spl66_48 | ~ spl66_62 | spl66_149), inference(avatar_contradiction_clause, [], [f3394])).
fof(f3394, plain, ($false | (~ spl66_48 | ~ spl66_62 | spl66_149)), inference(subsumption_resolution, [], [f3393, f803])).
fof(f3393, plain, (~ (e3 = op(e1, e0)) | (~ spl66_62 | spl66_149)), inference(forward_demodulation, [], [f1529, f863])).
fof(f1529, plain, (~ (e3 = op(op(e0, e0), e0)) | spl66_149), inference(avatar_component_clause, [], [f1527])).
fof(f1527, plain, (spl66_149 <=> (e3 = op(op(e0, e0), e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_149])])).
fof(f3387, plain, (~ spl66_26 | spl66_68), inference(avatar_contradiction_clause, [], [f3386])).
fof(f3386, plain, ($false | (~ spl66_26 | spl66_68)), inference(subsumption_resolution, [], [f3384, f710])).
fof(f710, plain, ((e1 = op(e2, e1)) | ~ spl66_26), inference(avatar_component_clause, [], [f708])).
fof(f3384, plain, (~ (e1 = op(e2, e1)) | (~ spl66_26 | spl66_68)), inference(backward_demodulation, [], [f924, f710])).
fof(f924, plain, (~ (e1 = op(e2, op(e2, e1))) | spl66_68), inference(avatar_component_clause, [], [f923])).
fof(f923, plain, (spl66_68 <=> (e1 = op(e2, op(e2, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl66_68])])).
fof(f3370, plain, (~ spl66_43 | ~ spl66_62 | spl66_150), inference(avatar_contradiction_clause, [], [f3369])).
fof(f3369, plain, ($false | (~ spl66_43 | ~ spl66_62 | spl66_150)), inference(subsumption_resolution, [], [f3362, f782])).
fof(f3362, plain, (~ (e2 = op(e1, e1)) | (~ spl66_62 | spl66_150)), inference(backward_demodulation, [], [f1534, f863])).
fof(f1534, plain, (~ (e2 = op(op(e0, e0), op(e0, e0))) | spl66_150), inference(avatar_component_clause, [], [f1532])).
fof(f1532, plain, (spl66_150 <=> (e2 = op(op(e0, e0), op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl66_150])])).
fof(f3368, plain, (spl66_57 | ~ spl66_62 | ~ spl66_79), inference(avatar_split_clause, [], [f3361, f982, f861, f840])).
fof(f3361, plain, ((e0 = op(e0, e1)) | (~ spl66_62 | ~ spl66_79)), inference(backward_demodulation, [], [f984, f863])).
fof(f3349, plain, (~ spl66_60 | ~ spl66_12), inference(avatar_split_clause, [], [f3348, f648, f852])).
fof(f3348, plain, (~ (e3 = op(e0, e1)) | ~ spl66_12), inference(forward_demodulation, [], [f243, f650])).
fof(f3347, plain, (~ spl66_59 | ~ spl66_43), inference(avatar_split_clause, [], [f3346, f780, f848])).
fof(f3346, plain, (~ (e2 = op(e0, e1)) | ~ spl66_43), inference(forward_demodulation, [], [f241, f782])).
fof(f3341, plain, (~ spl66_64 | ~ spl66_48), inference(avatar_split_clause, [], [f3340, f801, f869])).
fof(f869, plain, (spl66_64 <=> (op(e0, e0) = e3)), introduced(avatar_definition, [new_symbols(naming, [spl66_64])])).
fof(f3340, plain, (~ (op(e0, e0) = e3) | ~ spl66_48), inference(forward_demodulation, [], [f235, f803])).
fof(f235, plain, ~ (op(e0, e0) = op(e1, e0)), inference(cnf_transformation, [], [f3])).
fof(f3339, plain, (~ spl66_47 | ~ spl66_43), inference(avatar_split_clause, [], [f3338, f780, f797])).
fof(f3338, plain, (~ (e2 = op(e1, e0)) | ~ spl66_43), inference(forward_demodulation, [], [f265, f782])).
fof(f3337, plain, (spl66_33 | ~ spl66_48 | ~ spl66_74), inference(avatar_contradiction_clause, [], [f3336])).
fof(f3336, plain, ($false | (spl66_33 | ~ spl66_48 | ~ spl66_74)), inference(subsumption_resolution, [], [f3335, f739])).
fof(f739, plain, (~ (e0 = op(e1, e3)) | spl66_33), inference(avatar_component_clause, [], [f738])).
fof(f3335, plain, ((e0 = op(e1, e3)) | (~ spl66_48 | ~ spl66_74)), inference(forward_demodulation, [], [f957, f803])).
fof(f957, plain, ((e0 = op(e1, op(e1, e0))) | ~ spl66_74), inference(avatar_component_clause, [], [f955])).
fof(f955, plain, (spl66_74 <=> (e0 = op(e1, op(e1, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl66_74])])).
fof(f3329, plain, (~ spl66_27 | ~ spl66_43), inference(avatar_split_clause, [], [f3328, f780, f712])).
fof(f3328, plain, (~ (e2 = op(e2, e1)) | ~ spl66_43), inference(forward_demodulation, [], [f244, f782])).
fof(f244, plain, ~ (op(e1, e1) = op(e2, e1)), inference(cnf_transformation, [], [f3])).
fof(f3326, plain, (~ spl66_14 | ~ spl66_6), inference(avatar_split_clause, [], [f3325, f623, f657])).
fof(f3325, plain, (~ (e1 = op(e3, e0)) | ~ spl66_6), inference(forward_demodulation, [], [f278, f625])).
fof(f625, plain, ((e1 = op(e3, e2)) | ~ spl66_6), inference(avatar_component_clause, [], [f623])).
fof(f278, plain, ~ (op(e3, e0) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f3299, plain, (~ spl66_18 | ~ spl66_30), inference(avatar_split_clause, [], [f3297, f725, f674])).
fof(f3297, plain, (~ (e1 = op(e2, e3)) | ~ spl66_30), inference(backward_demodulation, [], [f273, f727])).
fof(f3298, plain, (~ spl66_14 | ~ spl66_30), inference(avatar_split_clause, [], [f3294, f725, f657])).
fof(f3294, plain, (~ (e1 = op(e3, e0)) | ~ spl66_30), inference(backward_demodulation, [], [f240, f727])).
fof(f240, plain, ~ (op(e2, e0) = op(e3, e0)), inference(cnf_transformation, [], [f3])).
fof(f3288, plain, (~ spl66_33 | ~ spl66_37), inference(avatar_split_clause, [], [f3284, f755, f738])).
fof(f3284, plain, (~ (e0 = op(e1, e3)) | ~ spl66_37), inference(backward_demodulation, [], [f270, f757])).
fof(f757, plain, ((e0 = op(e1, e2)) | ~ spl66_37), inference(avatar_component_clause, [], [f755])).
fof(f3287, plain, (~ spl66_5 | ~ spl66_37), inference(avatar_split_clause, [], [f3283, f755, f619])).
fof(f3283, plain, (~ (e0 = op(e3, e2)) | ~ spl66_37), inference(backward_demodulation, [], [f251, f757])).
fof(f3271, plain, (~ spl66_47 | ~ spl66_48), inference(avatar_contradiction_clause, [], [f3270])).
fof(f3270, plain, ($false | (~ spl66_47 | ~ spl66_48)), inference(subsumption_resolution, [], [f3269, f288])).
fof(f288, plain, ~ (e2 = e3), inference(cnf_transformation, [], [f4])).
fof(f3269, plain, ((e2 = e3) | (~ spl66_47 | ~ spl66_48)), inference(forward_demodulation, [], [f803, f799])).
fof(f3268, plain, (~ spl66_10 | ~ spl66_58), inference(avatar_split_clause, [], [f3266, f844, f640])).
fof(f3266, plain, (~ (e1 = op(e3, e1)) | ~ spl66_58), inference(backward_demodulation, [], [f243, f846])).
fof(f3257, plain, (~ spl66_57 | ~ spl66_61), inference(avatar_split_clause, [], [f3248, f857, f840])).
fof(f857, plain, (spl66_61 <=> (e0 = op(e0, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_61])])).
fof(f3248, plain, (~ (e0 = op(e0, e1)) | ~ spl66_61), inference(backward_demodulation, [], [f259, f859])).
fof(f859, plain, ((e0 = op(e0, e0)) | ~ spl66_61), inference(avatar_component_clause, [], [f857])).
fof(f3256, plain, (~ spl66_29 | ~ spl66_61), inference(avatar_split_clause, [], [f3246, f857, f721])).
fof(f3246, plain, (~ (e0 = op(e2, e0)) | ~ spl66_61), inference(backward_demodulation, [], [f236, f859])).
fof(f3245, plain, (spl66_37 | ~ spl66_47 | ~ spl66_74), inference(avatar_split_clause, [], [f3244, f955, f797, f755])).
fof(f3244, plain, ((e0 = op(e1, e2)) | (~ spl66_47 | ~ spl66_74)), inference(forward_demodulation, [], [f957, f799])).
fof(f3240, plain, (~ spl66_5 | ~ spl66_24 | spl66_163), inference(avatar_split_clause, [], [f3239, f1596, f699, f619])).
fof(f3239, plain, (~ (e0 = op(e3, e2)) | (~ spl66_24 | spl66_163)), inference(forward_demodulation, [], [f1598, f701])).
fof(f1598, plain, (~ (e0 = op(op(e2, e2), e2)) | spl66_163), inference(avatar_component_clause, [], [f1596])).
fof(f3229, plain, (~ spl66_17 | ~ spl66_3 | spl66_169), inference(avatar_split_clause, [], [f3169, f1626, f610, f670])).
fof(f3169, plain, (~ (e0 = op(e2, e3)) | (~ spl66_3 | spl66_169)), inference(backward_demodulation, [], [f1628, f612])).
fof(f3221, plain, (~ spl66_10 | spl66_145), inference(avatar_contradiction_clause, [], [f3220])).
fof(f3220, plain, ($false | (~ spl66_10 | spl66_145)), inference(subsumption_resolution, [], [f3219, f642])).
fof(f3219, plain, (~ (e1 = op(e3, e1)) | (~ spl66_10 | spl66_145)), inference(backward_demodulation, [], [f1505, f642])).
fof(f1505, plain, (~ (e1 = op(e3, op(e3, e1))) | spl66_145), inference(avatar_component_clause, [], [f1504])).
fof(f3188, plain, (~ spl66_63 | ~ spl66_47), inference(avatar_split_clause, [], [f3187, f797, f865])).
fof(f3187, plain, (~ (op(e0, e0) = e2) | ~ spl66_47), inference(forward_demodulation, [], [f235, f799])).
fof(f3184, plain, (~ spl66_32 | ~ spl66_24), inference(avatar_split_clause, [], [f3183, f699, f733])).
fof(f3183, plain, (~ (e3 = op(e2, e0)) | ~ spl66_24), inference(forward_demodulation, [], [f272, f701])).
fof(f3176, plain, (~ spl66_15 | ~ spl66_3), inference(avatar_split_clause, [], [f3175, f610, f661])).
fof(f3175, plain, (~ (e2 = op(e3, e0)) | ~ spl66_3), inference(forward_demodulation, [], [f279, f612])).
fof(f3157, plain, (~ spl66_8 | ~ spl66_12), inference(avatar_split_clause, [], [f3154, f648, f631])).
fof(f3154, plain, (~ (e3 = op(e3, e2)) | ~ spl66_12), inference(backward_demodulation, [], [f280, f650])).
fof(f3146, plain, (~ spl66_2 | ~ spl66_18), inference(avatar_split_clause, [], [f3142, f674, f606])).
fof(f3142, plain, (~ (e1 = op(e3, e3)) | ~ spl66_18), inference(backward_demodulation, [], [f258, f676])).
fof(f676, plain, ((e1 = op(e2, e3)) | ~ spl66_18), inference(avatar_component_clause, [], [f674])).
fof(f258, plain, ~ (op(e2, e3) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f3140, plain, (~ spl66_2 | ~ spl66_24 | spl66_162), inference(avatar_split_clause, [], [f3137, f1592, f699, f606])).
fof(f1592, plain, (spl66_162 <=> (e1 = op(op(e2, e2), op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl66_162])])).
fof(f3137, plain, (~ (e1 = op(e3, e3)) | (~ spl66_24 | spl66_162)), inference(backward_demodulation, [], [f1594, f701])).
fof(f1594, plain, (~ (e1 = op(op(e2, e2), op(e2, e2))) | spl66_162), inference(avatar_component_clause, [], [f1592])).
fof(f3139, plain, (~ spl66_8 | ~ spl66_24), inference(avatar_split_clause, [], [f3133, f699, f631])).
fof(f3133, plain, (~ (e3 = op(e3, e2)) | ~ spl66_24), inference(backward_demodulation, [], [f252, f701])).
fof(f252, plain, ~ (op(e2, e2) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f3132, plain, (~ spl66_25 | ~ spl66_27), inference(avatar_contradiction_clause, [], [f3131])).
fof(f3131, plain, ($false | (~ spl66_25 | ~ spl66_27)), inference(subsumption_resolution, [], [f3130, f284])).
fof(f3130, plain, ((e0 = e2) | (~ spl66_25 | ~ spl66_27)), inference(backward_demodulation, [], [f714, f706])).
fof(f706, plain, ((e0 = op(e2, e1)) | ~ spl66_25), inference(avatar_component_clause, [], [f704])).
fof(f704, plain, (spl66_25 <=> (e0 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_25])])).
fof(f3127, plain, (~ spl66_23 | ~ spl66_27), inference(avatar_split_clause, [], [f3124, f712, f695])).
fof(f3124, plain, (~ (e2 = op(e2, e2)) | ~ spl66_27), inference(backward_demodulation, [], [f274, f714])).
fof(f3120, plain, (~ spl66_17 | ~ spl66_33), inference(avatar_split_clause, [], [f3114, f738, f670])).
fof(f3114, plain, (~ (e0 = op(e2, e3)) | ~ spl66_33), inference(backward_demodulation, [], [f256, f740])).
fof(f3106, plain, (~ spl66_35 | ~ spl66_47), inference(avatar_split_clause, [], [f3102, f797, f746])).
fof(f3102, plain, (~ (e2 = op(e1, e3)) | ~ spl66_47), inference(backward_demodulation, [], [f267, f799])).
fof(f3105, plain, (~ spl66_15 | ~ spl66_47), inference(avatar_split_clause, [], [f3100, f797, f661])).
fof(f3100, plain, (~ (e2 = op(e3, e0)) | ~ spl66_47), inference(backward_demodulation, [], [f239, f799])).
fof(f3086, plain, (~ spl66_59 | ~ spl66_55), inference(avatar_split_clause, [], [f3085, f831, f848])).
fof(f3085, plain, (~ (e2 = op(e0, e1)) | ~ spl66_55), inference(forward_demodulation, [], [f262, f833])).
fof(f833, plain, ((e2 = op(e0, e2)) | ~ spl66_55), inference(avatar_component_clause, [], [f831])).
fof(f3079, plain, (~ spl66_23 | ~ spl66_55), inference(avatar_split_clause, [], [f3078, f831, f695])).
fof(f3078, plain, (~ (e2 = op(e2, e2)) | ~ spl66_55), inference(forward_demodulation, [], [f248, f833])).
fof(f3066, plain, (~ spl66_16 | ~ spl66_32), inference(avatar_split_clause, [], [f3065, f733, f665])).
fof(f3065, plain, (~ (e3 = op(e3, e0)) | ~ spl66_32), inference(forward_demodulation, [], [f240, f735])).
fof(f3064, plain, (~ spl66_13 | ~ spl66_61), inference(avatar_split_clause, [], [f3063, f857, f653])).
fof(f3063, plain, (~ (e0 = op(e3, e0)) | ~ spl66_61), inference(forward_demodulation, [], [f237, f859])).
fof(f3043, plain, (spl66_17 | ~ spl66_32 | ~ spl66_69), inference(avatar_split_clause, [], [f3041, f928, f733, f670])).
fof(f928, plain, (spl66_69 <=> (e0 = op(e2, op(e2, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl66_69])])).
fof(f3041, plain, ((e0 = op(e2, e3)) | (~ spl66_32 | ~ spl66_69)), inference(backward_demodulation, [], [f930, f735])).
fof(f930, plain, ((e0 = op(e2, op(e2, e0))) | ~ spl66_69), inference(avatar_component_clause, [], [f928])).
fof(f3042, plain, (~ spl66_20 | ~ spl66_32), inference(avatar_split_clause, [], [f3040, f733, f682])).
fof(f682, plain, (spl66_20 <=> (e3 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_20])])).
fof(f3040, plain, (~ (e3 = op(e2, e3)) | ~ spl66_32), inference(backward_demodulation, [], [f273, f735])).
fof(f3032, plain, (~ spl66_21 | ~ spl66_37), inference(avatar_split_clause, [], [f3026, f755, f687])).
fof(f3026, plain, (~ (e0 = op(e2, e2)) | ~ spl66_37), inference(backward_demodulation, [], [f250, f757])).
fof(f250, plain, ~ (op(e1, e2) = op(e2, e2)), inference(cnf_transformation, [], [f3])).
fof(f3025, plain, (~ spl66_20 | ~ spl66_52), inference(avatar_split_clause, [], [f3024, f818, f682])).
fof(f3024, plain, (~ (e3 = op(e2, e3)) | ~ spl66_52), inference(backward_demodulation, [], [f254, f820])).
fof(f3022, plain, (~ spl66_54 | ~ spl66_55), inference(avatar_contradiction_clause, [], [f3021])).
fof(f3021, plain, ($false | (~ spl66_54 | ~ spl66_55)), inference(subsumption_resolution, [], [f3020, f286])).
fof(f286, plain, ~ (e1 = e2), inference(cnf_transformation, [], [f4])).
fof(f3020, plain, ((e1 = e2) | (~ spl66_54 | ~ spl66_55)), inference(forward_demodulation, [], [f833, f829])).
fof(f2994, plain, (~ spl66_34 | ~ spl66_46), inference(avatar_split_clause, [], [f2993, f793, f742])).
fof(f2993, plain, (~ (e1 = op(e1, e3)) | ~ spl66_46), inference(forward_demodulation, [], [f267, f795])).
fof(f2983, plain, (~ spl66_13 | spl66_146), inference(avatar_contradiction_clause, [], [f2982])).
fof(f2982, plain, ($false | (~ spl66_13 | spl66_146)), inference(subsumption_resolution, [], [f2980, f655])).
fof(f2980, plain, (~ (e0 = op(e3, e0)) | (~ spl66_13 | spl66_146)), inference(backward_demodulation, [], [f1510, f655])).
fof(f1510, plain, (~ (e0 = op(e3, op(e3, e0))) | spl66_146), inference(avatar_component_clause, [], [f1509])).
fof(f2975, plain, (~ spl66_21 | ~ spl66_64 | spl66_164), inference(avatar_contradiction_clause, [], [f2974])).
fof(f2974, plain, ($false | (~ spl66_21 | ~ spl66_64 | spl66_164)), inference(subsumption_resolution, [], [f2972, f871])).
fof(f871, plain, ((op(e0, e0) = e3) | ~ spl66_64), inference(avatar_component_clause, [], [f869])).
fof(f2972, plain, (~ (op(e0, e0) = e3) | (~ spl66_21 | spl66_164)), inference(backward_demodulation, [], [f1604, f689])).
fof(f1604, plain, (~ (e3 = op(op(e2, e2), op(e2, e2))) | spl66_164), inference(avatar_component_clause, [], [f1602])).
fof(f1602, plain, (spl66_164 <=> (e3 = op(op(e2, e2), op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl66_164])])).
fof(f2973, plain, (~ spl66_5 | ~ spl66_21), inference(avatar_split_clause, [], [f2966, f687, f619])).
fof(f2966, plain, (~ (e0 = op(e3, e2)) | ~ spl66_21), inference(backward_demodulation, [], [f252, f689])).
fof(f2961, plain, (spl66_21 | ~ spl66_31 | ~ spl66_69), inference(avatar_split_clause, [], [f2958, f928, f729, f687])).
fof(f2958, plain, ((e0 = op(e2, e2)) | (~ spl66_31 | ~ spl66_69)), inference(backward_demodulation, [], [f930, f731])).
fof(f2960, plain, (~ spl66_23 | ~ spl66_31), inference(avatar_split_clause, [], [f2956, f729, f695])).
fof(f2956, plain, (~ (e2 = op(e2, e2)) | ~ spl66_31), inference(backward_demodulation, [], [f272, f731])).
fof(f2953, plain, (~ spl66_39 | spl66_72), inference(avatar_contradiction_clause, [], [f2952])).
fof(f2952, plain, ($false | (~ spl66_39 | spl66_72)), inference(subsumption_resolution, [], [f2950, f765])).
fof(f2950, plain, (~ (e2 = op(e1, e2)) | (~ spl66_39 | spl66_72)), inference(backward_demodulation, [], [f946, f765])).
fof(f2951, plain, (~ spl66_23 | ~ spl66_39), inference(avatar_split_clause, [], [f2947, f763, f695])).
fof(f2947, plain, (~ (e2 = op(e2, e2)) | ~ spl66_39), inference(backward_demodulation, [], [f250, f765])).
fof(f2946, plain, (~ spl66_38 | ~ spl66_46), inference(avatar_split_clause, [], [f2941, f793, f759])).
fof(f2941, plain, (~ (e1 = op(e1, e2)) | ~ spl66_46), inference(backward_demodulation, [], [f266, f795])).
fof(f2945, plain, (~ spl66_30 | ~ spl66_46), inference(avatar_split_clause, [], [f2939, f793, f725])).
fof(f2939, plain, (~ (e1 = op(e2, e0)) | ~ spl66_46), inference(backward_demodulation, [], [f238, f795])).
fof(f2937, plain, (~ spl66_38 | ~ spl66_54), inference(avatar_split_clause, [], [f2931, f827, f759])).
fof(f2931, plain, (~ (e1 = op(e1, e2)) | ~ spl66_54), inference(backward_demodulation, [], [f247, f829])).
fof(f2923, plain, (~ spl66_2 | ~ spl66_64 | spl66_147), inference(avatar_contradiction_clause, [], [f2922])).
fof(f2922, plain, ($false | (~ spl66_2 | ~ spl66_64 | spl66_147)), inference(subsumption_resolution, [], [f2916, f608])).
fof(f2916, plain, (~ (e1 = op(e3, e3)) | (~ spl66_64 | spl66_147)), inference(backward_demodulation, [], [f1520, f871])).
fof(f1520, plain, (~ (e1 = op(op(e0, e0), op(e0, e0))) | spl66_147), inference(avatar_component_clause, [], [f1518])).
fof(f1518, plain, (spl66_147 <=> (e1 = op(op(e0, e0), op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl66_147])])).
fof(f2921, plain, (~ spl66_56 | ~ spl66_64), inference(avatar_split_clause, [], [f2914, f869, f835])).
fof(f2914, plain, (~ (e3 = op(e0, e2)) | ~ spl66_64), inference(backward_demodulation, [], [f260, f871])).
fof(f2920, plain, (~ spl66_16 | ~ spl66_64), inference(avatar_split_clause, [], [f2912, f869, f665])).
fof(f2912, plain, (~ (e3 = op(e3, e0)) | ~ spl66_64), inference(backward_demodulation, [], [f237, f871])).
fof(f2896, plain, (~ spl66_59 | ~ spl66_11), inference(avatar_split_clause, [], [f2895, f644, f848])).
fof(f2895, plain, (~ (e2 = op(e0, e1)) | ~ spl66_11), inference(forward_demodulation, [], [f243, f646])).
fof(f2885, plain, (~ spl66_5 | spl66_15 | ~ spl66_144), inference(avatar_contradiction_clause, [], [f2884])).
fof(f2884, plain, ($false | (~ spl66_5 | spl66_15 | ~ spl66_144)), inference(subsumption_resolution, [], [f2882, f662])).
fof(f662, plain, (~ (e2 = op(e3, e0)) | spl66_15), inference(avatar_component_clause, [], [f661])).
fof(f2882, plain, ((e2 = op(e3, e0)) | (~ spl66_5 | ~ spl66_144)), inference(backward_demodulation, [], [f1501, f621])).
fof(f621, plain, ((e0 = op(e3, e2)) | ~ spl66_5), inference(avatar_component_clause, [], [f619])).
fof(f2879, plain, (~ spl66_9 | ~ spl66_11), inference(avatar_contradiction_clause, [], [f2878])).
fof(f2878, plain, ($false | (~ spl66_9 | ~ spl66_11)), inference(subsumption_resolution, [], [f2877, f284])).
fof(f2877, plain, ((e0 = e2) | (~ spl66_9 | ~ spl66_11)), inference(forward_demodulation, [], [f646, f638])).
fof(f638, plain, ((e0 = op(e3, e1)) | ~ spl66_9), inference(avatar_component_clause, [], [f636])).
fof(f2863, plain, (spl66_22 | ~ spl66_27 | ~ spl66_68), inference(avatar_contradiction_clause, [], [f2862])).
fof(f2862, plain, ($false | (spl66_22 | ~ spl66_27 | ~ spl66_68)), inference(subsumption_resolution, [], [f2859, f692])).
fof(f692, plain, (~ (e1 = op(e2, e2)) | spl66_22), inference(avatar_component_clause, [], [f691])).
fof(f2859, plain, ((e1 = op(e2, e2)) | (~ spl66_27 | ~ spl66_68)), inference(backward_demodulation, [], [f925, f714])).
fof(f925, plain, ((e1 = op(e2, op(e2, e1))) | ~ spl66_68), inference(avatar_component_clause, [], [f923])).
fof(f2858, plain, (spl66_25 | ~ spl66_30 | ~ spl66_69), inference(avatar_contradiction_clause, [], [f2857])).
fof(f2857, plain, ($false | (spl66_25 | ~ spl66_30 | ~ spl66_69)), inference(subsumption_resolution, [], [f2855, f705])).
fof(f705, plain, (~ (e0 = op(e2, e1)) | spl66_25), inference(avatar_component_clause, [], [f704])).
fof(f2855, plain, ((e0 = op(e2, e1)) | (~ spl66_30 | ~ spl66_69)), inference(backward_demodulation, [], [f930, f727])).
fof(f2844, plain, (~ spl66_31 | ~ spl66_47), inference(avatar_split_clause, [], [f2842, f797, f729])).
fof(f2842, plain, (~ (e2 = op(e2, e0)) | ~ spl66_47), inference(backward_demodulation, [], [f238, f799])).
fof(f2840, plain, (~ spl66_35 | ~ spl66_51), inference(avatar_split_clause, [], [f2839, f814, f746])).
fof(f2839, plain, (~ (e2 = op(e1, e3)) | ~ spl66_51), inference(backward_demodulation, [], [f253, f816])).
fof(f2828, plain, (~ spl66_49 | ~ spl66_61), inference(avatar_split_clause, [], [f2818, f857, f806])).
fof(f806, plain, (spl66_49 <=> (e0 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_49])])).
fof(f2818, plain, (~ (e0 = op(e0, e3)) | ~ spl66_61), inference(backward_demodulation, [], [f261, f859])).
fof(f261, plain, ~ (op(e0, e0) = op(e0, e3)), inference(cnf_transformation, [], [f3])).
fof(f2827, plain, (~ spl66_45 | ~ spl66_61), inference(avatar_split_clause, [], [f2815, f857, f789])).
fof(f2815, plain, (~ (e0 = op(e1, e0)) | ~ spl66_61), inference(backward_demodulation, [], [f235, f859])).
fof(f2814, plain, (spl66_51 | ~ spl66_56 | ~ spl66_77), inference(avatar_split_clause, [], [f2813, f972, f835, f814])).
fof(f972, plain, (spl66_77 <=> (e2 = op(e0, op(e0, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl66_77])])).
fof(f2813, plain, ((e2 = op(e0, e3)) | (~ spl66_56 | ~ spl66_77)), inference(forward_demodulation, [], [f974, f837])).
fof(f974, plain, ((e2 = op(e0, op(e0, e2))) | ~ spl66_77), inference(avatar_component_clause, [], [f972])).
fof(f2803, plain, (~ spl66_36 | ~ spl66_20), inference(avatar_split_clause, [], [f2696, f682, f750])).
fof(f750, plain, (spl66_36 <=> (e3 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_36])])).
fof(f2696, plain, (~ (e3 = op(e1, e3)) | ~ spl66_20), inference(forward_demodulation, [], [f256, f684])).
fof(f684, plain, ((e3 = op(e2, e3)) | ~ spl66_20), inference(avatar_component_clause, [], [f682])).
fof(f2800, plain, (~ spl66_33 | ~ spl66_2 | spl66_169), inference(avatar_split_clause, [], [f2779, f1626, f606, f738])).
fof(f2779, plain, (~ (e0 = op(e1, e3)) | (~ spl66_2 | spl66_169)), inference(backward_demodulation, [], [f1628, f608])).
fof(f2799, plain, (~ spl66_22 | ~ spl66_38), inference(avatar_split_clause, [], [f2798, f759, f691])).
fof(f2798, plain, (~ (e1 = op(e2, e2)) | ~ spl66_38), inference(forward_demodulation, [], [f250, f761])).
fof(f2795, plain, (~ spl66_24 | ~ spl66_20), inference(avatar_split_clause, [], [f2702, f682, f699])).
fof(f2702, plain, (~ (e3 = op(e2, e2)) | ~ spl66_20), inference(forward_demodulation, [], [f276, f684])).
fof(f2790, plain, (~ spl66_5 | ~ spl66_9), inference(avatar_split_clause, [], [f2789, f636, f619])).
fof(f2789, plain, (~ (e0 = op(e3, e2)) | ~ spl66_9), inference(forward_demodulation, [], [f280, f638])).
fof(f2765, plain, (~ spl66_38 | spl66_43 | ~ spl66_72), inference(avatar_contradiction_clause, [], [f2764])).
fof(f2764, plain, ($false | (~ spl66_38 | spl66_43 | ~ spl66_72)), inference(subsumption_resolution, [], [f2762, f781])).
fof(f781, plain, (~ (e2 = op(e1, e1)) | spl66_43), inference(avatar_component_clause, [], [f780])).
fof(f2762, plain, ((e2 = op(e1, e1)) | (~ spl66_38 | ~ spl66_72)), inference(backward_demodulation, [], [f947, f761])).
fof(f947, plain, ((e2 = op(e1, op(e1, e2))) | ~ spl66_72), inference(avatar_component_clause, [], [f945])).
fof(f2760, plain, (~ spl66_45 | spl66_74), inference(avatar_contradiction_clause, [], [f2759])).
fof(f2759, plain, ($false | (~ spl66_45 | spl66_74)), inference(subsumption_resolution, [], [f2756, f791])).
fof(f2756, plain, (~ (e0 = op(e1, e0)) | (~ spl66_45 | spl66_74)), inference(backward_demodulation, [], [f956, f791])).
fof(f956, plain, (~ (e0 = op(e1, op(e1, e0))) | spl66_74), inference(avatar_component_clause, [], [f955])).
fof(f2758, plain, (~ spl66_29 | ~ spl66_45), inference(avatar_split_clause, [], [f2755, f789, f721])).
fof(f2755, plain, (~ (e0 = op(e2, e0)) | ~ spl66_45), inference(backward_demodulation, [], [f238, f791])).
fof(f2750, plain, (~ spl66_1 | ~ spl66_49), inference(avatar_split_clause, [], [f2748, f806, f602])).
fof(f2748, plain, (~ (e0 = op(e3, e3)) | ~ spl66_49), inference(backward_demodulation, [], [f255, f808])).
fof(f808, plain, ((e0 = op(e0, e3)) | ~ spl66_49), inference(avatar_component_clause, [], [f806])).
fof(f2738, plain, (~ spl66_46 | ~ spl66_62), inference(avatar_split_clause, [], [f2727, f861, f793])).
fof(f2727, plain, (~ (e1 = op(e1, e0)) | ~ spl66_62), inference(backward_demodulation, [], [f235, f863])).
fof(f2720, plain, (~ spl66_62 | ~ spl66_21 | spl66_162), inference(avatar_split_clause, [], [f2558, f1592, f687, f861])).
fof(f2558, plain, (~ (op(e0, e0) = e1) | (~ spl66_21 | spl66_162)), inference(forward_demodulation, [], [f1594, f689])).
fof(f2706, plain, (spl66_40 | ~ spl66_35 | ~ spl66_71), inference(avatar_split_clause, [], [f2654, f940, f746, f767])).
fof(f940, plain, (spl66_71 <=> (e3 = op(e1, op(e1, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl66_71])])).
fof(f2654, plain, ((e3 = op(e1, e2)) | (~ spl66_35 | ~ spl66_71)), inference(backward_demodulation, [], [f942, f748])).
fof(f942, plain, ((e3 = op(e1, op(e1, e3))) | ~ spl66_71), inference(avatar_component_clause, [], [f940])).
fof(f2705, plain, (spl66_31 | ~ spl66_21 | ~ spl66_67), inference(avatar_split_clause, [], [f2604, f918, f687, f729])).
fof(f2604, plain, ((e2 = op(e2, e0)) | (~ spl66_21 | ~ spl66_67)), inference(forward_demodulation, [], [f920, f689])).
fof(f2692, plain, (~ spl66_1 | ~ spl66_51 | spl66_167), inference(avatar_contradiction_clause, [], [f2691])).
fof(f2691, plain, ($false | (~ spl66_1 | ~ spl66_51 | spl66_167)), inference(subsumption_resolution, [], [f2686, f816])).
fof(f2686, plain, (~ (e2 = op(e0, e3)) | (~ spl66_1 | spl66_167)), inference(backward_demodulation, [], [f1619, f604])).
fof(f1619, plain, (~ (e2 = op(op(e3, e3), e3)) | spl66_167), inference(avatar_component_clause, [], [f1617])).
fof(f2690, plain, (~ spl66_1 | ~ spl66_15 | ~ spl66_143), inference(avatar_contradiction_clause, [], [f2689])).
fof(f2689, plain, ($false | (~ spl66_1 | ~ spl66_15 | ~ spl66_143)), inference(subsumption_resolution, [], [f2688, f288])).
fof(f2688, plain, ((e2 = e3) | (~ spl66_1 | ~ spl66_15 | ~ spl66_143)), inference(forward_demodulation, [], [f2684, f663])).
fof(f2684, plain, ((e3 = op(e3, e0)) | (~ spl66_1 | ~ spl66_143)), inference(backward_demodulation, [], [f1496, f604])).
fof(f2676, plain, (~ spl66_15 | ~ spl66_16), inference(avatar_contradiction_clause, [], [f2675])).
fof(f2675, plain, ($false | (~ spl66_15 | ~ spl66_16)), inference(subsumption_resolution, [], [f2674, f288])).
fof(f2674, plain, ((e2 = e3) | (~ spl66_15 | ~ spl66_16)), inference(backward_demodulation, [], [f667, f663])).
fof(f2650, plain, (~ spl66_37 | ~ spl66_46 | ~ spl66_72), inference(avatar_contradiction_clause, [], [f2649])).
fof(f2649, plain, ($false | (~ spl66_37 | ~ spl66_46 | ~ spl66_72)), inference(subsumption_resolution, [], [f2648, f286])).
fof(f2648, plain, ((e1 = e2) | (~ spl66_37 | ~ spl66_46 | ~ spl66_72)), inference(forward_demodulation, [], [f2644, f795])).
fof(f2644, plain, ((e2 = op(e1, e0)) | (~ spl66_37 | ~ spl66_72)), inference(backward_demodulation, [], [f947, f757])).
fof(f2642, plain, (~ spl66_9 | ~ spl66_44 | spl66_157), inference(avatar_contradiction_clause, [], [f2641])).
fof(f2641, plain, ($false | (~ spl66_9 | ~ spl66_44 | spl66_157)), inference(subsumption_resolution, [], [f2635, f638])).
fof(f2635, plain, (~ (e0 = op(e3, e1)) | (~ spl66_44 | spl66_157)), inference(backward_demodulation, [], [f1568, f786])).
fof(f1568, plain, (~ (e0 = op(op(e1, e1), e1)) | spl66_157), inference(avatar_component_clause, [], [f1566])).
fof(f1566, plain, (spl66_157 <=> (e0 = op(op(e1, e1), e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_157])])).
fof(f2627, plain, (~ spl66_3 | ~ spl66_51), inference(avatar_split_clause, [], [f2625, f814, f610])).
fof(f2625, plain, (~ (e2 = op(e3, e3)) | ~ spl66_51), inference(backward_demodulation, [], [f255, f816])).
fof(f2623, plain, (~ spl66_40 | ~ spl66_56), inference(avatar_split_clause, [], [f2620, f835, f767])).
fof(f2620, plain, (~ (e3 = op(e1, e2)) | ~ spl66_56), inference(backward_demodulation, [], [f247, f837])).
fof(f2602, plain, (spl66_1 | ~ spl66_16 | ~ spl66_146), inference(avatar_split_clause, [], [f2601, f1509, f665, f602])).
fof(f2601, plain, ((e0 = op(e3, e3)) | (~ spl66_16 | ~ spl66_146)), inference(forward_demodulation, [], [f1511, f667])).
fof(f2589, plain, (~ spl66_61 | spl66_79), inference(avatar_contradiction_clause, [], [f2588])).
fof(f2588, plain, ($false | (~ spl66_61 | spl66_79)), inference(subsumption_resolution, [], [f2587, f859])).
fof(f2587, plain, (~ (e0 = op(e0, e0)) | (~ spl66_61 | spl66_79)), inference(forward_demodulation, [], [f983, f859])).
fof(f983, plain, (~ (e0 = op(e0, op(e0, e0))) | spl66_79), inference(avatar_component_clause, [], [f982])).
fof(f2572, plain, (~ spl66_25 | ~ spl66_21), inference(avatar_split_clause, [], [f2571, f687, f704])).
fof(f2571, plain, (~ (e0 = op(e2, e1)) | ~ spl66_21), inference(forward_demodulation, [], [f274, f689])).
fof(f2551, plain, (~ spl66_1 | ~ spl66_9), inference(avatar_split_clause, [], [f2550, f636, f602])).
fof(f2550, plain, (~ (e0 = op(e3, e3)) | ~ spl66_9), inference(forward_demodulation, [], [f281, f638])).
fof(f2548, plain, (~ spl66_9 | spl66_14 | ~ spl66_145), inference(avatar_contradiction_clause, [], [f2547])).
fof(f2547, plain, ($false | (~ spl66_9 | spl66_14 | ~ spl66_145)), inference(subsumption_resolution, [], [f2546, f658])).
fof(f658, plain, (~ (e1 = op(e3, e0)) | spl66_14), inference(avatar_component_clause, [], [f657])).
fof(f2546, plain, ((e1 = op(e3, e0)) | (~ spl66_9 | ~ spl66_145)), inference(forward_demodulation, [], [f1506, f638])).
fof(f2545, plain, (~ spl66_25 | ~ spl66_9), inference(avatar_split_clause, [], [f2544, f636, f704])).
fof(f2544, plain, (~ (e0 = op(e2, e1)) | ~ spl66_9), inference(forward_demodulation, [], [f246, f638])).
fof(f2541, plain, (~ spl66_3 | ~ spl66_18 | spl66_166), inference(avatar_contradiction_clause, [], [f2540])).
fof(f2540, plain, ($false | (~ spl66_3 | ~ spl66_18 | spl66_166)), inference(subsumption_resolution, [], [f2537, f676])).
fof(f2537, plain, (~ (e1 = op(e2, e3)) | (~ spl66_3 | spl66_166)), inference(backward_demodulation, [], [f1614, f612])).
fof(f1614, plain, (~ (e1 = op(op(e3, e3), e3)) | spl66_166), inference(avatar_component_clause, [], [f1612])).
fof(f1612, plain, (spl66_166 <=> (e1 = op(op(e3, e3), e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_166])])).
fof(f2526, plain, (~ spl66_21 | ~ spl66_23), inference(avatar_contradiction_clause, [], [f2525])).
fof(f2525, plain, ($false | (~ spl66_21 | ~ spl66_23)), inference(subsumption_resolution, [], [f2524, f284])).
fof(f2524, plain, ((e0 = e2) | (~ spl66_21 | ~ spl66_23)), inference(backward_demodulation, [], [f697, f689])).
fof(f697, plain, ((e2 = op(e2, e2)) | ~ spl66_23), inference(avatar_component_clause, [], [f695])).
fof(f2520, plain, (~ spl66_23 | ~ spl66_161), inference(avatar_contradiction_clause, [], [f2519])).
fof(f2519, plain, ($false | (~ spl66_23 | ~ spl66_161)), inference(subsumption_resolution, [], [f2518, f288])).
fof(f2518, plain, ((e2 = e3) | (~ spl66_23 | ~ spl66_161)), inference(forward_demodulation, [], [f2510, f697])).
fof(f2510, plain, ((e3 = op(e2, e2)) | (~ spl66_23 | ~ spl66_161)), inference(backward_demodulation, [], [f1588, f697])).
fof(f1588, plain, ((e3 = op(op(e2, e2), e2)) | ~ spl66_161), inference(avatar_component_clause, [], [f1587])).
fof(f2515, plain, (~ spl66_23 | spl66_67), inference(avatar_contradiction_clause, [], [f2514])).
fof(f2514, plain, ($false | (~ spl66_23 | spl66_67)), inference(subsumption_resolution, [], [f2507, f697])).
fof(f2507, plain, (~ (e2 = op(e2, e2)) | (~ spl66_23 | spl66_67)), inference(backward_demodulation, [], [f919, f697])).
fof(f919, plain, (~ (e2 = op(e2, op(e2, e2))) | spl66_67), inference(avatar_component_clause, [], [f918])).
fof(f2488, plain, (~ spl66_61 | ~ spl66_63), inference(avatar_contradiction_clause, [], [f2487])).
fof(f2487, plain, ($false | (~ spl66_61 | ~ spl66_63)), inference(subsumption_resolution, [], [f2485, f284])).
fof(f2485, plain, ((e0 = e2) | (~ spl66_61 | ~ spl66_63)), inference(backward_demodulation, [], [f867, f859])).
fof(f2480, plain, (~ spl66_33 | ~ spl66_40 | ~ spl66_72), inference(avatar_contradiction_clause, [], [f2479])).
fof(f2479, plain, ($false | (~ spl66_33 | ~ spl66_40 | ~ spl66_72)), inference(subsumption_resolution, [], [f2478, f284])).
fof(f2478, plain, ((e0 = e2) | (~ spl66_33 | ~ spl66_40 | ~ spl66_72)), inference(forward_demodulation, [], [f2477, f740])).
fof(f2477, plain, ((e2 = op(e1, e3)) | (~ spl66_40 | ~ spl66_72)), inference(forward_demodulation, [], [f947, f769])).
fof(f2476, plain, (spl66_38 | ~ spl66_43 | ~ spl66_73), inference(avatar_contradiction_clause, [], [f2475])).
fof(f2475, plain, ($false | (spl66_38 | ~ spl66_43 | ~ spl66_73)), inference(subsumption_resolution, [], [f2474, f760])).
fof(f760, plain, (~ (e1 = op(e1, e2)) | spl66_38), inference(avatar_component_clause, [], [f759])).
fof(f2474, plain, ((e1 = op(e1, e2)) | (~ spl66_43 | ~ spl66_73)), inference(forward_demodulation, [], [f952, f782])).
fof(f2473, plain, (spl66_41 | ~ spl66_46 | ~ spl66_74), inference(avatar_contradiction_clause, [], [f2472])).
fof(f2472, plain, ($false | (spl66_41 | ~ spl66_46 | ~ spl66_74)), inference(subsumption_resolution, [], [f2471, f773])).
fof(f773, plain, (~ (e0 = op(e1, e1)) | spl66_41), inference(avatar_component_clause, [], [f772])).
fof(f2471, plain, ((e0 = op(e1, e1)) | (~ spl66_46 | ~ spl66_74)), inference(forward_demodulation, [], [f957, f795])).
fof(f2468, plain, (~ spl66_22 | ~ spl66_63 | spl66_147), inference(avatar_split_clause, [], [f2467, f1518, f865, f691])).
fof(f2467, plain, (~ (e1 = op(e2, e2)) | (~ spl66_63 | spl66_147)), inference(forward_demodulation, [], [f1520, f867])).
fof(f2457, plain, (spl66_59 | ~ spl66_54 | ~ spl66_77), inference(avatar_split_clause, [], [f2409, f972, f827, f848])).
fof(f2409, plain, ((e2 = op(e0, e1)) | (~ spl66_54 | ~ spl66_77)), inference(backward_demodulation, [], [f974, f829])).
fof(f2435, plain, (~ spl66_28 | ~ spl66_43 | spl66_155), inference(avatar_split_clause, [], [f2377, f1557, f780, f716])).
fof(f2377, plain, (~ (e3 = op(e2, e1)) | (~ spl66_43 | spl66_155)), inference(forward_demodulation, [], [f1559, f782])).
fof(f2434, plain, (~ spl66_25 | ~ spl66_43 | spl66_157), inference(avatar_split_clause, [], [f2382, f1566, f780, f704])).
fof(f2382, plain, (~ (e0 = op(e2, e1)) | (~ spl66_43 | spl66_157)), inference(forward_demodulation, [], [f1568, f782])).
fof(f2420, plain, (~ spl66_37 | ~ spl66_40), inference(avatar_contradiction_clause, [], [f2419])).
fof(f2419, plain, ($false | (~ spl66_37 | ~ spl66_40)), inference(subsumption_resolution, [], [f2418, f285])).
fof(f285, plain, ~ (e0 = e3), inference(cnf_transformation, [], [f4])).
fof(f2418, plain, ((e0 = e3) | (~ spl66_37 | ~ spl66_40)), inference(forward_demodulation, [], [f769, f757])).
fof(f2416, plain, (~ spl66_50 | ~ spl66_52), inference(avatar_contradiction_clause, [], [f2415])).
fof(f2415, plain, ($false | (~ spl66_50 | ~ spl66_52)), inference(subsumption_resolution, [], [f2414, f287])).
fof(f2414, plain, ((e1 = e3) | (~ spl66_50 | ~ spl66_52)), inference(forward_demodulation, [], [f820, f812])).
fof(f2393, plain, (~ spl66_1 | ~ spl66_63 | spl66_170), inference(avatar_contradiction_clause, [], [f2392])).
fof(f2392, plain, ($false | (~ spl66_1 | ~ spl66_63 | spl66_170)), inference(subsumption_resolution, [], [f2391, f867])).
fof(f2391, plain, (~ (op(e0, e0) = e2) | (~ spl66_1 | spl66_170)), inference(forward_demodulation, [], [f1634, f604])).
fof(f2355, plain, (~ spl66_1 | ~ spl66_50 | spl66_166), inference(avatar_contradiction_clause, [], [f2354])).
fof(f2354, plain, ($false | (~ spl66_1 | ~ spl66_50 | spl66_166)), inference(subsumption_resolution, [], [f2347, f812])).
fof(f2347, plain, (~ (e1 = op(e0, e3)) | (~ spl66_1 | spl66_166)), inference(backward_demodulation, [], [f1614, f604])).
fof(f2337, plain, (~ spl66_6 | ~ spl66_10), inference(avatar_split_clause, [], [f2335, f640, f623])).
fof(f2335, plain, (~ (e1 = op(e3, e2)) | ~ spl66_10), inference(backward_demodulation, [], [f280, f642])).
fof(f2327, plain, (~ spl66_43 | ~ spl66_44), inference(avatar_contradiction_clause, [], [f2326])).
fof(f2326, plain, ($false | (~ spl66_43 | ~ spl66_44)), inference(subsumption_resolution, [], [f2325, f288])).
fof(f2325, plain, ((e2 = e3) | (~ spl66_43 | ~ spl66_44)), inference(backward_demodulation, [], [f786, f782])).
fof(f2324, plain, (~ spl66_2 | ~ spl66_50), inference(avatar_split_clause, [], [f2320, f810, f606])).
fof(f2320, plain, (~ (e1 = op(e3, e3)) | ~ spl66_50), inference(backward_demodulation, [], [f255, f812])).
fof(f2316, plain, (~ spl66_57 | spl66_62 | ~ spl66_78), inference(avatar_contradiction_clause, [], [f2315])).
fof(f2315, plain, ($false | (~ spl66_57 | spl66_62 | ~ spl66_78)), inference(subsumption_resolution, [], [f2313, f862])).
fof(f862, plain, (~ (op(e0, e0) = e1) | spl66_62), inference(avatar_component_clause, [], [f861])).
fof(f2313, plain, ((op(e0, e0) = e1) | (~ spl66_57 | ~ spl66_78)), inference(backward_demodulation, [], [f979, f842])).
fof(f979, plain, ((e1 = op(e0, op(e0, e1))) | ~ spl66_78), inference(avatar_component_clause, [], [f977])).
fof(f977, plain, (spl66_78 <=> (e1 = op(e0, op(e0, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl66_78])])).
fof(f2314, plain, (~ spl66_49 | ~ spl66_57), inference(avatar_split_clause, [], [f2312, f840, f806])).
fof(f2312, plain, (~ (e0 = op(e0, e3)) | ~ spl66_57), inference(backward_demodulation, [], [f263, f842])).
fof(f2305, plain, (~ spl66_22 | ~ spl66_44 | spl66_164), inference(avatar_contradiction_clause, [], [f2304])).
fof(f2304, plain, ($false | (~ spl66_22 | ~ spl66_44 | spl66_164)), inference(subsumption_resolution, [], [f2303, f786])).
fof(f2303, plain, (~ (e3 = op(e1, e1)) | (~ spl66_22 | spl66_164)), inference(forward_demodulation, [], [f1604, f693])).
fof(f2298, plain, (~ spl66_60 | ~ spl66_56), inference(avatar_split_clause, [], [f2297, f835, f852])).
fof(f2297, plain, (~ (e3 = op(e0, e1)) | ~ spl66_56), inference(forward_demodulation, [], [f262, f837])).
fof(f2276, plain, (~ spl66_36 | ~ spl66_44), inference(avatar_contradiction_clause, [], [f2275])).
fof(f2275, plain, ($false | (~ spl66_36 | ~ spl66_44)), inference(subsumption_resolution, [], [f2092, f786])).
fof(f2092, plain, (~ (e3 = op(e1, e1)) | ~ spl66_36), inference(forward_demodulation, [], [f269, f752])).
fof(f752, plain, ((e3 = op(e1, e3)) | ~ spl66_36), inference(avatar_component_clause, [], [f750])).
fof(f2272, plain, (~ spl66_11 | ~ spl66_44 | spl66_154), inference(avatar_split_clause, [], [f2271, f1552, f784, f644])).
fof(f1552, plain, (spl66_154 <=> (e2 = op(op(e1, e1), e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_154])])).
fof(f2271, plain, (~ (e2 = op(e3, e1)) | (~ spl66_44 | spl66_154)), inference(forward_demodulation, [], [f1554, f786])).
fof(f1554, plain, (~ (e2 = op(op(e1, e1), e1)) | spl66_154), inference(avatar_component_clause, [], [f1552])).
fof(f2270, plain, (~ spl66_40 | ~ spl66_36), inference(avatar_split_clause, [], [f2094, f750, f767])).
fof(f2094, plain, (~ (e3 = op(e1, e2)) | ~ spl66_36), inference(forward_demodulation, [], [f270, f752])).
fof(f2267, plain, (~ spl66_29 | spl66_69), inference(avatar_contradiction_clause, [], [f2266])).
fof(f2266, plain, ($false | (~ spl66_29 | spl66_69)), inference(subsumption_resolution, [], [f2265, f723])).
fof(f2265, plain, (~ (e0 = op(e2, e0)) | (~ spl66_29 | spl66_69)), inference(forward_demodulation, [], [f929, f723])).
fof(f929, plain, (~ (e0 = op(e2, op(e2, e0))) | spl66_69), inference(avatar_component_clause, [], [f928])).
fof(f2251, plain, (~ spl66_4 | ~ spl66_36), inference(avatar_split_clause, [], [f2250, f750, f614])).
fof(f2250, plain, (~ (e3 = op(e3, e3)) | ~ spl66_36), inference(forward_demodulation, [], [f257, f752])).
fof(f2244, plain, (spl66_6 | ~ spl66_11 | ~ spl66_145), inference(avatar_split_clause, [], [f2241, f1504, f644, f623])).
fof(f2241, plain, ((e1 = op(e3, e2)) | (~ spl66_11 | ~ spl66_145)), inference(backward_demodulation, [], [f1506, f646])).
fof(f2240, plain, (~ spl66_13 | ~ spl66_16), inference(avatar_contradiction_clause, [], [f2239])).
fof(f2239, plain, ($false | (~ spl66_13 | ~ spl66_16)), inference(subsumption_resolution, [], [f2238, f285])).
fof(f2238, plain, ((e0 = e3) | (~ spl66_13 | ~ spl66_16)), inference(forward_demodulation, [], [f667, f655])).
fof(f2226, plain, (~ spl66_17 | ~ spl66_49), inference(avatar_split_clause, [], [f2221, f806, f670])).
fof(f2221, plain, (~ (e0 = op(e2, e3)) | ~ spl66_49), inference(backward_demodulation, [], [f254, f808])).
fof(f2225, plain, (~ spl66_49 | ~ spl66_63 | ~ spl66_76), inference(avatar_contradiction_clause, [], [f2224])).
fof(f2224, plain, ($false | (~ spl66_49 | ~ spl66_63 | ~ spl66_76)), inference(subsumption_resolution, [], [f2223, f288])).
fof(f2223, plain, ((e2 = e3) | (~ spl66_49 | ~ spl66_63 | ~ spl66_76)), inference(forward_demodulation, [], [f2220, f867])).
fof(f2220, plain, ((op(e0, e0) = e3) | (~ spl66_49 | ~ spl66_76)), inference(backward_demodulation, [], [f969, f808])).
fof(f969, plain, ((e3 = op(e0, op(e0, e3))) | ~ spl66_76), inference(avatar_component_clause, [], [f967])).
fof(f2216, plain, (~ spl66_53 | ~ spl66_56), inference(avatar_contradiction_clause, [], [f2215])).
fof(f2215, plain, ($false | (~ spl66_53 | ~ spl66_56)), inference(subsumption_resolution, [], [f2214, f285])).
fof(f2214, plain, ((e0 = e3) | (~ spl66_53 | ~ spl66_56)), inference(forward_demodulation, [], [f837, f825])).
fof(f2205, plain, (~ spl66_59 | ~ spl66_41 | spl66_154), inference(avatar_split_clause, [], [f2204, f1552, f772, f848])).
fof(f2204, plain, (~ (e2 = op(e0, e1)) | (~ spl66_41 | spl66_154)), inference(forward_demodulation, [], [f1554, f774])).
fof(f2202, plain, (~ spl66_49 | ~ spl66_53), inference(avatar_split_clause, [], [f2103, f823, f806])).
fof(f2103, plain, (~ (e0 = op(e0, e3)) | ~ spl66_53), inference(forward_demodulation, [], [f264, f825])).
fof(f2189, plain, (~ spl66_32 | ~ spl66_63 | spl66_149), inference(avatar_contradiction_clause, [], [f2188])).
fof(f2188, plain, ($false | (~ spl66_32 | ~ spl66_63 | spl66_149)), inference(subsumption_resolution, [], [f2118, f735])).
fof(f2118, plain, (~ (e3 = op(e2, e0)) | (~ spl66_63 | spl66_149)), inference(forward_demodulation, [], [f1529, f867])).
fof(f2179, plain, (~ spl66_6 | ~ spl66_22), inference(avatar_split_clause, [], [f2178, f691, f623])).
fof(f2178, plain, (~ (e1 = op(e3, e2)) | ~ spl66_22), inference(forward_demodulation, [], [f252, f693])).
fof(f2176, plain, (~ spl66_6 | ~ spl66_2), inference(avatar_split_clause, [], [f2175, f606, f623])).
fof(f2175, plain, (~ (e1 = op(e3, e2)) | ~ spl66_2), inference(forward_demodulation, [], [f282, f608])).
fof(f2174, plain, (~ spl66_1 | ~ spl66_2), inference(avatar_contradiction_clause, [], [f2173])).
fof(f2173, plain, ($false | (~ spl66_1 | ~ spl66_2)), inference(subsumption_resolution, [], [f2172, f283])).
fof(f2172, plain, ((e0 = e1) | (~ spl66_1 | ~ spl66_2)), inference(backward_demodulation, [], [f608, f604])).
fof(f2171, plain, (~ spl66_2 | ~ spl66_3), inference(avatar_contradiction_clause, [], [f2170])).
fof(f2170, plain, ($false | (~ spl66_2 | ~ spl66_3)), inference(subsumption_resolution, [], [f2169, f286])).
fof(f2169, plain, ((e1 = e2) | (~ spl66_2 | ~ spl66_3)), inference(forward_demodulation, [], [f612, f608])).
fof(f2165, plain, (~ spl66_31 | ~ spl66_32), inference(avatar_contradiction_clause, [], [f2164])).
fof(f2164, plain, ($false | (~ spl66_31 | ~ spl66_32)), inference(subsumption_resolution, [], [f2163, f288])).
fof(f2163, plain, ((e2 = e3) | (~ spl66_31 | ~ spl66_32)), inference(forward_demodulation, [], [f735, f731])).
fof(f2161, plain, (~ spl66_53 | ~ spl66_54), inference(avatar_contradiction_clause, [], [f2160])).
fof(f2160, plain, ($false | (~ spl66_53 | ~ spl66_54)), inference(subsumption_resolution, [], [f2159, f283])).
fof(f2159, plain, ((e0 = e1) | (~ spl66_53 | ~ spl66_54)), inference(forward_demodulation, [], [f829, f825])).
fof(f2157, plain, (~ spl66_53 | ~ spl66_59 | ~ spl66_78), inference(avatar_contradiction_clause, [], [f2156])).
fof(f2156, plain, ($false | (~ spl66_53 | ~ spl66_59 | ~ spl66_78)), inference(subsumption_resolution, [], [f2155, f283])).
fof(f2155, plain, ((e0 = e1) | (~ spl66_53 | ~ spl66_59 | ~ spl66_78)), inference(forward_demodulation, [], [f2152, f825])).
fof(f2152, plain, ((e1 = op(e0, e2)) | (~ spl66_59 | ~ spl66_78)), inference(backward_demodulation, [], [f979, f850])).
fof(f2145, plain, (~ spl66_2 | spl66_12 | ~ spl66_143), inference(avatar_contradiction_clause, [], [f2144])).
fof(f2144, plain, ($false | (~ spl66_2 | spl66_12 | ~ spl66_143)), inference(subsumption_resolution, [], [f2143, f649])).
fof(f649, plain, (~ (e3 = op(e3, e1)) | spl66_12), inference(avatar_component_clause, [], [f648])).
fof(f2143, plain, ((e3 = op(e3, e1)) | (~ spl66_2 | ~ spl66_143)), inference(forward_demodulation, [], [f1496, f608])).
fof(f2126, plain, (~ spl66_51 | ~ spl66_63), inference(avatar_split_clause, [], [f2125, f865, f814])).
fof(f2125, plain, (~ (e2 = op(e0, e3)) | ~ spl66_63), inference(forward_demodulation, [], [f261, f867])).
fof(f2121, plain, (~ spl66_30 | ~ spl66_63 | spl66_151), inference(avatar_split_clause, [], [f2120, f1536, f865, f725])).
fof(f1536, plain, (spl66_151 <=> (e1 = op(op(e0, e0), e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_151])])).
fof(f2120, plain, (~ (e1 = op(e2, e0)) | (~ spl66_63 | spl66_151)), inference(forward_demodulation, [], [f1538, f867])).
fof(f1538, plain, (~ (e1 = op(op(e0, e0), e0)) | spl66_151), inference(avatar_component_clause, [], [f1536])).
fof(f2089, plain, (~ spl66_52 | ~ spl66_36), inference(avatar_split_clause, [], [f2088, f750, f818])).
fof(f2088, plain, (~ (e3 = op(e0, e3)) | ~ spl66_36), inference(forward_demodulation, [], [f253, f752])).
fof(f2086, plain, (~ spl66_35 | ~ spl66_2 | spl66_167), inference(avatar_split_clause, [], [f2076, f1617, f606, f746])).
fof(f2076, plain, (~ (e2 = op(e1, e3)) | (~ spl66_2 | spl66_167)), inference(backward_demodulation, [], [f1619, f608])).
fof(f2082, plain, (~ spl66_1 | ~ spl66_17), inference(avatar_split_clause, [], [f2062, f670, f602])).
fof(f2062, plain, (~ (e0 = op(e3, e3)) | ~ spl66_17), inference(backward_demodulation, [], [f258, f672])).
fof(f2070, plain, (~ spl66_3 | ~ spl66_11), inference(avatar_split_clause, [], [f2069, f644, f610])).
fof(f2069, plain, (~ (e2 = op(e3, e3)) | ~ spl66_11), inference(backward_demodulation, [], [f281, f646])).
fof(f2053, plain, (~ spl66_34 | ~ spl66_36), inference(avatar_contradiction_clause, [], [f2052])).
fof(f2052, plain, ($false | (~ spl66_34 | ~ spl66_36)), inference(subsumption_resolution, [], [f2051, f287])).
fof(f2051, plain, ((e1 = e3) | (~ spl66_34 | ~ spl66_36)), inference(forward_demodulation, [], [f752, f744])).
fof(f2047, plain, (~ spl66_51 | ~ spl66_53 | ~ spl66_76), inference(avatar_contradiction_clause, [], [f2046])).
fof(f2046, plain, ($false | (~ spl66_51 | ~ spl66_53 | ~ spl66_76)), inference(subsumption_resolution, [], [f2045, f285])).
fof(f2045, plain, ((e0 = e3) | (~ spl66_51 | ~ spl66_53 | ~ spl66_76)), inference(forward_demodulation, [], [f2041, f825])).
fof(f2041, plain, ((e3 = op(e0, e2)) | (~ spl66_51 | ~ spl66_76)), inference(backward_demodulation, [], [f969, f816])).
fof(f2040, plain, (~ spl66_53 | ~ spl66_55), inference(avatar_contradiction_clause, [], [f2039])).
fof(f2039, plain, ($false | (~ spl66_53 | ~ spl66_55)), inference(subsumption_resolution, [], [f2038, f284])).
fof(f2038, plain, ((e0 = e2) | (~ spl66_53 | ~ spl66_55)), inference(backward_demodulation, [], [f833, f825])).
fof(f2036, plain, (~ spl66_57 | ~ spl66_58), inference(avatar_contradiction_clause, [], [f2035])).
fof(f2035, plain, ($false | (~ spl66_57 | ~ spl66_58)), inference(subsumption_resolution, [], [f2034, f283])).
fof(f2034, plain, ((e0 = e1) | (~ spl66_57 | ~ spl66_58)), inference(backward_demodulation, [], [f846, f842])).
fof(f2033, plain, (~ spl66_63 | ~ spl66_64), inference(avatar_contradiction_clause, [], [f2032])).
fof(f2032, plain, ($false | (~ spl66_63 | ~ spl66_64)), inference(subsumption_resolution, [], [f2031, f288])).
fof(f2031, plain, ((e2 = e3) | (~ spl66_63 | ~ spl66_64)), inference(backward_demodulation, [], [f871, f867])).
fof(f2028, plain, (spl66_49 | ~ spl66_64 | ~ spl66_79), inference(avatar_split_clause, [], [f2027, f982, f869, f806])).
fof(f2027, plain, ((e0 = op(e0, e3)) | (~ spl66_64 | ~ spl66_79)), inference(forward_demodulation, [], [f984, f871])).
fof(f2024, plain, (~ spl66_3 | ~ spl66_64 | spl66_150), inference(avatar_split_clause, [], [f2023, f1532, f869, f610])).
fof(f2023, plain, (~ (e2 = op(e3, e3)) | (~ spl66_64 | spl66_150)), inference(forward_demodulation, [], [f1534, f871])).
fof(f2017, plain, (~ spl66_51 | ~ spl66_55), inference(avatar_split_clause, [], [f2016, f831, f814])).
fof(f2016, plain, (~ (e2 = op(e0, e3)) | ~ spl66_55), inference(forward_demodulation, [], [f264, f833])).
fof(f2013, plain, (~ spl66_39 | ~ spl66_55), inference(avatar_split_clause, [], [f1930, f831, f763])).
fof(f1930, plain, (~ (e2 = op(e1, e2)) | ~ spl66_55), inference(forward_demodulation, [], [f247, f833])).
fof(f2002, plain, (~ spl66_3 | ~ spl66_21 | spl66_165), inference(avatar_contradiction_clause, [], [f2001])).
fof(f2001, plain, ($false | (~ spl66_3 | ~ spl66_21 | spl66_165)), inference(subsumption_resolution, [], [f1994, f689])).
fof(f1994, plain, (~ (e0 = op(e2, e2)) | (~ spl66_3 | spl66_165)), inference(backward_demodulation, [], [f1610, f612])).
fof(f1610, plain, (~ (e0 = op(op(e3, e3), op(e3, e3))) | spl66_165), inference(avatar_component_clause, [], [f1608])).
fof(f1608, plain, (spl66_165 <=> (e0 = op(op(e3, e3), op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl66_165])])).
fof(f2000, plain, (~ spl66_3 | ~ spl66_6 | ~ spl66_143), inference(avatar_contradiction_clause, [], [f1999])).
fof(f1999, plain, ($false | (~ spl66_3 | ~ spl66_6 | ~ spl66_143)), inference(subsumption_resolution, [], [f1998, f287])).
fof(f1998, plain, ((e1 = e3) | (~ spl66_3 | ~ spl66_6 | ~ spl66_143)), inference(forward_demodulation, [], [f1993, f625])).
fof(f1993, plain, ((e3 = op(e3, e2)) | (~ spl66_3 | ~ spl66_143)), inference(backward_demodulation, [], [f1496, f612])).
fof(f1991, plain, (~ spl66_6 | ~ spl66_8), inference(avatar_contradiction_clause, [], [f1990])).
fof(f1990, plain, ($false | (~ spl66_6 | ~ spl66_8)), inference(subsumption_resolution, [], [f1988, f287])).
fof(f1988, plain, ((e1 = e3) | (~ spl66_6 | ~ spl66_8)), inference(backward_demodulation, [], [f633, f625])).
fof(f1975, plain, (~ spl66_21 | ~ spl66_22), inference(avatar_contradiction_clause, [], [f1974])).
fof(f1974, plain, ($false | (~ spl66_21 | ~ spl66_22)), inference(subsumption_resolution, [], [f1973, f283])).
fof(f1973, plain, ((e0 = e1) | (~ spl66_21 | ~ spl66_22)), inference(backward_demodulation, [], [f693, f689])).
fof(f1961, plain, (~ spl66_41 | ~ spl66_43), inference(avatar_contradiction_clause, [], [f1960])).
fof(f1960, plain, ($false | (~ spl66_41 | ~ spl66_43)), inference(subsumption_resolution, [], [f1959, f284])).
fof(f1959, plain, ((e0 = e2) | (~ spl66_41 | ~ spl66_43)), inference(forward_demodulation, [], [f782, f774])).
fof(f1948, plain, (~ spl66_14 | ~ spl66_64 | spl66_151), inference(avatar_split_clause, [], [f1947, f1536, f869, f657])).
fof(f1947, plain, (~ (e1 = op(e3, e0)) | (~ spl66_64 | spl66_151)), inference(forward_demodulation, [], [f1538, f871])).
fof(f1942, plain, (~ spl66_41 | ~ spl66_64 | spl66_158), inference(avatar_contradiction_clause, [], [f1941])).
fof(f1941, plain, ($false | (~ spl66_41 | ~ spl66_64 | spl66_158)), inference(subsumption_resolution, [], [f1940, f871])).
fof(f1940, plain, (~ (op(e0, e0) = e3) | (~ spl66_41 | spl66_158)), inference(forward_demodulation, [], [f1574, f774])).
fof(f1932, plain, (~ spl66_37 | ~ spl66_41), inference(avatar_split_clause, [], [f1931, f772, f755])).
fof(f1931, plain, (~ (e0 = op(e1, e2)) | ~ spl66_41), inference(forward_demodulation, [], [f268, f774])).
fof(f1918, plain, (~ spl66_25 | ~ spl66_41), inference(avatar_split_clause, [], [f1917, f772, f704])).
fof(f1917, plain, (~ (e0 = op(e2, e1)) | ~ spl66_41), inference(forward_demodulation, [], [f244, f774])).
fof(f1913, plain, (~ spl66_18 | ~ spl66_22), inference(avatar_split_clause, [], [f1912, f691, f674])).
fof(f1912, plain, (~ (e1 = op(e2, e3)) | ~ spl66_22), inference(forward_demodulation, [], [f276, f693])).
fof(f1904, plain, (~ spl66_16 | ~ spl66_8), inference(avatar_split_clause, [], [f1903, f631, f665])).
fof(f1903, plain, (~ (e3 = op(e3, e0)) | ~ spl66_8), inference(forward_demodulation, [], [f278, f633])).
fof(f1893, plain, (~ spl66_2 | ~ spl66_41 | spl66_165), inference(avatar_contradiction_clause, [], [f1892])).
fof(f1892, plain, ($false | (~ spl66_2 | ~ spl66_41 | spl66_165)), inference(subsumption_resolution, [], [f1885, f774])).
fof(f1885, plain, (~ (e0 = op(e1, e1)) | (~ spl66_2 | spl66_165)), inference(backward_demodulation, [], [f1610, f608])).
fof(f1880, plain, (~ spl66_5 | ~ spl66_8), inference(avatar_contradiction_clause, [], [f1879])).
fof(f1879, plain, ($false | (~ spl66_5 | ~ spl66_8)), inference(subsumption_resolution, [], [f1878, f285])).
fof(f1878, plain, ((e0 = e3) | (~ spl66_5 | ~ spl66_8)), inference(forward_demodulation, [], [f633, f621])).
fof(f1873, plain, (~ spl66_37 | ~ spl66_22 | spl66_163), inference(avatar_split_clause, [], [f1867, f1596, f691, f755])).
fof(f1867, plain, (~ (e0 = op(e1, e2)) | (~ spl66_22 | spl66_163)), inference(backward_demodulation, [], [f1598, f693])).
fof(f1870, plain, (~ spl66_22 | spl66_27 | ~ spl66_67), inference(avatar_contradiction_clause, [], [f1869])).
fof(f1869, plain, ($false | (~ spl66_22 | spl66_27 | ~ spl66_67)), inference(subsumption_resolution, [], [f1865, f713])).
fof(f713, plain, (~ (e2 = op(e2, e1)) | spl66_27), inference(avatar_component_clause, [], [f712])).
fof(f1865, plain, ((e2 = op(e2, e1)) | (~ spl66_22 | ~ spl66_67)), inference(backward_demodulation, [], [f920, f693])).
fof(f1863, plain, (~ spl66_29 | ~ spl66_31), inference(avatar_contradiction_clause, [], [f1862])).
fof(f1862, plain, ($false | (~ spl66_29 | ~ spl66_31)), inference(subsumption_resolution, [], [f1861, f284])).
fof(f1861, plain, ((e0 = e2) | (~ spl66_29 | ~ spl66_31)), inference(backward_demodulation, [], [f731, f723])).
fof(f1859, plain, (~ spl66_35 | ~ spl66_36), inference(avatar_contradiction_clause, [], [f1858])).
fof(f1858, plain, ($false | (~ spl66_35 | ~ spl66_36)), inference(subsumption_resolution, [], [f1857, f288])).
fof(f1857, plain, ((e2 = e3) | (~ spl66_35 | ~ spl66_36)), inference(backward_demodulation, [], [f752, f748])).
fof(f1841, plain, (~ spl66_5 | ~ spl66_11 | ~ spl66_145), inference(avatar_contradiction_clause, [], [f1840])).
fof(f1840, plain, ($false | (~ spl66_5 | ~ spl66_11 | ~ spl66_145)), inference(subsumption_resolution, [], [f1839, f283])).
fof(f1839, plain, ((e0 = e1) | (~ spl66_5 | ~ spl66_11 | ~ spl66_145)), inference(forward_demodulation, [], [f1838, f621])).
fof(f1838, plain, ((e1 = op(e3, e2)) | (~ spl66_11 | ~ spl66_145)), inference(forward_demodulation, [], [f1506, f646])).
fof(f1821, plain, (~ spl66_45 | ~ spl66_41), inference(avatar_split_clause, [], [f1820, f772, f789])).
fof(f1820, plain, (~ (e0 = op(e1, e0)) | ~ spl66_41), inference(forward_demodulation, [], [f265, f774])).
fof(f1817, plain, (~ spl66_48 | ~ spl66_36), inference(avatar_split_clause, [], [f1816, f750, f801])).
fof(f1816, plain, (~ (e3 = op(e1, e0)) | ~ spl66_36), inference(forward_demodulation, [], [f267, f752])).
fof(f1803, plain, (~ spl66_63 | ~ spl66_41 | spl66_156), inference(avatar_split_clause, [], [f1802, f1562, f772, f865])).
fof(f1802, plain, (~ (op(e0, e0) = e2) | (~ spl66_41 | spl66_156)), inference(forward_demodulation, [], [f1564, f774])).
fof(f1783, plain, (~ spl66_1 | ~ spl66_62 | spl66_168), inference(avatar_contradiction_clause, [], [f1782])).
fof(f1782, plain, ($false | (~ spl66_1 | ~ spl66_62 | spl66_168)), inference(subsumption_resolution, [], [f1780, f863])).
fof(f1780, plain, (~ (op(e0, e0) = e1) | (~ spl66_1 | spl66_168)), inference(backward_demodulation, [], [f1624, f604])).
fof(f1775, plain, (~ spl66_1 | ~ spl66_5), inference(avatar_split_clause, [], [f1774, f619, f602])).
fof(f1774, plain, (~ (e0 = op(e3, e3)) | ~ spl66_5), inference(backward_demodulation, [], [f282, f621])).
fof(f1767, plain, (~ spl66_14 | ~ spl66_16), inference(avatar_contradiction_clause, [], [f1766])).
fof(f1766, plain, ($false | (~ spl66_14 | ~ spl66_16)), inference(subsumption_resolution, [], [f1765, f287])).
fof(f1765, plain, ((e1 = e3) | (~ spl66_14 | ~ spl66_16)), inference(backward_demodulation, [], [f667, f659])).
fof(f1764, plain, (~ spl66_4 | ~ spl66_16), inference(avatar_split_clause, [], [f1761, f665, f614])).
fof(f1761, plain, (~ (e3 = op(e3, e3)) | ~ spl66_16), inference(backward_demodulation, [], [f279, f667])).
fof(f1755, plain, (~ spl66_18 | ~ spl66_19), inference(avatar_contradiction_clause, [], [f1754])).
fof(f1754, plain, ($false | (~ spl66_18 | ~ spl66_19)), inference(subsumption_resolution, [], [f1753, f286])).
fof(f1753, plain, ((e1 = e2) | (~ spl66_18 | ~ spl66_19)), inference(backward_demodulation, [], [f680, f676])).
fof(f1752, plain, (~ spl66_3 | ~ spl66_19), inference(avatar_split_clause, [], [f1750, f678, f610])).
fof(f1750, plain, (~ (e2 = op(e3, e3)) | ~ spl66_19), inference(backward_demodulation, [], [f258, f680])).
fof(f1748, plain, (~ spl66_1 | ~ spl66_24 | spl66_159), inference(avatar_split_clause, [], [f1743, f1578, f699, f602])).
fof(f1743, plain, (~ (e0 = op(e3, e3)) | (~ spl66_24 | spl66_159)), inference(backward_demodulation, [], [f1580, f701])).
fof(f1739, plain, (spl66_18 | ~ spl66_28 | ~ spl66_68), inference(avatar_split_clause, [], [f1735, f923, f716, f674])).
fof(f1735, plain, ((e1 = op(e2, e3)) | (~ spl66_28 | ~ spl66_68)), inference(backward_demodulation, [], [f925, f718])).
fof(f1727, plain, (~ spl66_15 | ~ spl66_31), inference(avatar_split_clause, [], [f1722, f729, f661])).
fof(f1722, plain, (~ (e2 = op(e3, e0)) | ~ spl66_31), inference(backward_demodulation, [], [f240, f731])).
fof(f1717, plain, (~ spl66_38 | ~ spl66_39), inference(avatar_contradiction_clause, [], [f1716])).
fof(f1716, plain, ($false | (~ spl66_38 | ~ spl66_39)), inference(subsumption_resolution, [], [f1715, f286])).
fof(f1715, plain, ((e1 = e2) | (~ spl66_38 | ~ spl66_39)), inference(backward_demodulation, [], [f765, f761])).
fof(f1708, plain, (~ spl66_41 | ~ spl66_42), inference(avatar_contradiction_clause, [], [f1707])).
fof(f1707, plain, ($false | (~ spl66_41 | ~ spl66_42)), inference(subsumption_resolution, [], [f1706, f283])).
fof(f1706, plain, ((e0 = e1) | (~ spl66_41 | ~ spl66_42)), inference(backward_demodulation, [], [f778, f774])).
fof(f778, plain, ((e1 = op(e1, e1)) | ~ spl66_42), inference(avatar_component_clause, [], [f776])).
fof(f776, plain, (spl66_42 <=> (e1 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_42])])).
fof(f1703, plain, (~ spl66_34 | ~ spl66_42), inference(avatar_split_clause, [], [f1695, f776, f742])).
fof(f1695, plain, (~ (e1 = op(e1, e3)) | ~ spl66_42), inference(backward_demodulation, [], [f269, f778])).
fof(f1702, plain, (~ spl66_38 | ~ spl66_42), inference(avatar_split_clause, [], [f1694, f776, f759])).
fof(f1694, plain, (~ (e1 = op(e1, e2)) | ~ spl66_42), inference(backward_demodulation, [], [f268, f778])).
fof(f1701, plain, (~ spl66_10 | ~ spl66_42), inference(avatar_split_clause, [], [f1693, f776, f640])).
fof(f1693, plain, (~ (e1 = op(e3, e1)) | ~ spl66_42), inference(backward_demodulation, [], [f245, f778])).
fof(f1678, plain, (~ spl66_33 | ~ spl66_49), inference(avatar_split_clause, [], [f1675, f806, f738])).
fof(f1675, plain, (~ (e0 = op(e1, e3)) | ~ spl66_49), inference(backward_demodulation, [], [f253, f808])).
fof(f1656, plain, (~ spl66_62 | ~ spl66_64), inference(avatar_contradiction_clause, [], [f1655])).
fof(f1655, plain, ($false | (~ spl66_62 | ~ spl66_64)), inference(subsumption_resolution, [], [f1654, f287])).
fof(f1654, plain, ((e1 = e3) | (~ spl66_62 | ~ spl66_64)), inference(backward_demodulation, [], [f871, f863])).
fof(f1651, plain, (~ spl66_15 | ~ spl66_64 | spl66_148), inference(avatar_split_clause, [], [f1643, f1522, f869, f661])).
fof(f1643, plain, (~ (e2 = op(e3, e0)) | (~ spl66_64 | spl66_148)), inference(backward_demodulation, [], [f1524, f871])).
fof(f1650, plain, (~ spl66_52 | ~ spl66_64), inference(avatar_split_clause, [], [f1642, f869, f818])).
fof(f1642, plain, (~ (e3 = op(e0, e3)) | ~ spl66_64), inference(backward_demodulation, [], [f261, f871])).
fof(f1648, plain, (~ spl66_60 | ~ spl66_64), inference(avatar_split_clause, [], [f1640, f869, f852])).
fof(f1640, plain, (~ (e3 = op(e0, e1)) | ~ spl66_64), inference(backward_demodulation, [], [f259, f871])).
fof(f1646, plain, (~ spl66_32 | ~ spl66_64), inference(avatar_split_clause, [], [f1638, f869, f733])).
fof(f1638, plain, (~ (e3 = op(e2, e0)) | ~ spl66_64), inference(backward_demodulation, [], [f236, f871])).
fof(f1636, plain, (~ spl66_170 | ~ spl66_166 | ~ spl66_1), inference(avatar_split_clause, [], [f600, f602, f1612, f1632])).
fof(f600, plain, (~ (e0 = op(e3, e3)) | ~ (e1 = op(op(e3, e3), e3)) | ~ (e2 = op(op(e3, e3), op(e3, e3)))), inference(cnf_transformation, [], [f53])).
fof(f53, plain, (~ (e0 = op(e3, e3)) | ~ (e1 = op(op(e3, e3), e3)) | ~ (e2 = op(op(e3, e3), op(e3, e3)))), inference(ennf_transformation, [], [f29])).
fof(f29, plain, ~ ((e0 = op(e3, e3)) & (e1 = op(op(e3, e3), e3)) & (e2 = op(op(e3, e3), op(e3, e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG154+1.p', ax29)).
fof(f1635, plain, (~ spl66_170 | ~ spl66_169 | ~ spl66_2), inference(avatar_split_clause, [], [f599, f606, f1626, f1632])).
fof(f599, plain, (~ (e1 = op(e3, e3)) | ~ (e0 = op(op(e3, e3), e3)) | ~ (e2 = op(op(e3, e3), op(e3, e3)))), inference(cnf_transformation, [], [f52])).
fof(f52, plain, (~ (e1 = op(e3, e3)) | ~ (e0 = op(op(e3, e3), e3)) | ~ (e2 = op(op(e3, e3), op(e3, e3)))), inference(ennf_transformation, [], [f28])).
fof(f28, plain, ~ ((e1 = op(e3, e3)) & (e0 = op(op(e3, e3), e3)) & (e2 = op(op(e3, e3), op(e3, e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG154+1.p', ax28)).
fof(f1630, plain, (~ spl66_168 | ~ spl66_167 | ~ spl66_1), inference(avatar_split_clause, [], [f598, f602, f1617, f1622])).
fof(f598, plain, (~ (e0 = op(e3, e3)) | ~ (e2 = op(op(e3, e3), e3)) | ~ (e1 = op(op(e3, e3), op(e3, e3)))), inference(cnf_transformation, [], [f51])).
fof(f51, plain, (~ (e0 = op(e3, e3)) | ~ (e2 = op(op(e3, e3), e3)) | ~ (e1 = op(op(e3, e3), op(e3, e3)))), inference(ennf_transformation, [], [f27])).
fof(f27, plain, ~ ((e0 = op(e3, e3)) & (e2 = op(op(e3, e3), e3)) & (e1 = op(op(e3, e3), op(e3, e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG154+1.p', ax27)).
fof(f1629, plain, (~ spl66_168 | ~ spl66_169 | ~ spl66_3), inference(avatar_split_clause, [], [f597, f610, f1626, f1622])).
fof(f597, plain, (~ (e2 = op(e3, e3)) | ~ (e0 = op(op(e3, e3), e3)) | ~ (e1 = op(op(e3, e3), op(e3, e3)))), inference(cnf_transformation, [], [f50])).
fof(f50, plain, (~ (e2 = op(e3, e3)) | ~ (e0 = op(op(e3, e3), e3)) | ~ (e1 = op(op(e3, e3), op(e3, e3)))), inference(ennf_transformation, [], [f26])).
fof(f26, plain, ~ ((e2 = op(e3, e3)) & (e0 = op(op(e3, e3), e3)) & (e1 = op(op(e3, e3), op(e3, e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG154+1.p', ax26)).
fof(f1620, plain, (~ spl66_165 | ~ spl66_167 | ~ spl66_2), inference(avatar_split_clause, [], [f596, f606, f1617, f1608])).
fof(f596, plain, (~ (e1 = op(e3, e3)) | ~ (e2 = op(op(e3, e3), e3)) | ~ (e0 = op(op(e3, e3), op(e3, e3)))), inference(cnf_transformation, [], [f49])).
fof(f49, plain, (~ (e1 = op(e3, e3)) | ~ (e2 = op(op(e3, e3), e3)) | ~ (e0 = op(op(e3, e3), op(e3, e3)))), inference(ennf_transformation, [], [f25])).
fof(f25, plain, ~ ((e1 = op(e3, e3)) & (e2 = op(op(e3, e3), e3)) & (e0 = op(op(e3, e3), op(e3, e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG154+1.p', ax25)).
fof(f1615, plain, (~ spl66_165 | ~ spl66_166 | ~ spl66_3), inference(avatar_split_clause, [], [f595, f610, f1612, f1608])).
fof(f595, plain, (~ (e2 = op(e3, e3)) | ~ (e1 = op(op(e3, e3), e3)) | ~ (e0 = op(op(e3, e3), op(e3, e3)))), inference(cnf_transformation, [], [f48])).
fof(f48, plain, (~ (e2 = op(e3, e3)) | ~ (e1 = op(op(e3, e3), e3)) | ~ (e0 = op(op(e3, e3), op(e3, e3)))), inference(ennf_transformation, [], [f24])).
fof(f24, plain, ~ ((e2 = op(e3, e3)) & (e1 = op(op(e3, e3), e3)) & (e0 = op(op(e3, e3), op(e3, e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG154+1.p', ax24)).
fof(f1606, plain, (~ spl66_164 | ~ spl66_160 | ~ spl66_21), inference(avatar_split_clause, [], [f594, f687, f1582, f1602])).
fof(f594, plain, (~ (e0 = op(e2, e2)) | ~ (e1 = op(op(e2, e2), e2)) | ~ (e3 = op(op(e2, e2), op(e2, e2)))), inference(cnf_transformation, [], [f47])).
fof(f47, plain, (~ (e0 = op(e2, e2)) | ~ (e1 = op(op(e2, e2), e2)) | ~ (e3 = op(op(e2, e2), op(e2, e2)))), inference(ennf_transformation, [], [f23])).
fof(f23, plain, ~ ((e0 = op(e2, e2)) & (e1 = op(op(e2, e2), e2)) & (e3 = op(op(e2, e2), op(e2, e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG154+1.p', ax23)).
fof(f1605, plain, (~ spl66_164 | ~ spl66_163 | ~ spl66_22), inference(avatar_split_clause, [], [f593, f691, f1596, f1602])).
fof(f593, plain, (~ (e1 = op(e2, e2)) | ~ (e0 = op(op(e2, e2), e2)) | ~ (e3 = op(op(e2, e2), op(e2, e2)))), inference(cnf_transformation, [], [f46])).
fof(f46, plain, (~ (e1 = op(e2, e2)) | ~ (e0 = op(op(e2, e2), e2)) | ~ (e3 = op(op(e2, e2), op(e2, e2)))), inference(ennf_transformation, [], [f22])).
fof(f22, plain, ~ ((e1 = op(e2, e2)) & (e0 = op(op(e2, e2), e2)) & (e3 = op(op(e2, e2), op(e2, e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG154+1.p', ax22)).
fof(f1600, plain, (~ spl66_162 | ~ spl66_161 | ~ spl66_21), inference(avatar_split_clause, [], [f592, f687, f1587, f1592])).
fof(f592, plain, (~ (e0 = op(e2, e2)) | ~ (e3 = op(op(e2, e2), e2)) | ~ (e1 = op(op(e2, e2), op(e2, e2)))), inference(cnf_transformation, [], [f45])).
fof(f45, plain, (~ (e0 = op(e2, e2)) | ~ (e3 = op(op(e2, e2), e2)) | ~ (e1 = op(op(e2, e2), op(e2, e2)))), inference(ennf_transformation, [], [f21])).
fof(f21, plain, ~ ((e0 = op(e2, e2)) & (e3 = op(op(e2, e2), e2)) & (e1 = op(op(e2, e2), op(e2, e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG154+1.p', ax21)).
fof(f1599, plain, (~ spl66_162 | ~ spl66_163 | ~ spl66_24), inference(avatar_split_clause, [], [f591, f699, f1596, f1592])).
fof(f591, plain, (~ (e3 = op(e2, e2)) | ~ (e0 = op(op(e2, e2), e2)) | ~ (e1 = op(op(e2, e2), op(e2, e2)))), inference(cnf_transformation, [], [f44])).
fof(f44, plain, (~ (e3 = op(e2, e2)) | ~ (e0 = op(op(e2, e2), e2)) | ~ (e1 = op(op(e2, e2), op(e2, e2)))), inference(ennf_transformation, [], [f20])).
fof(f20, plain, ~ ((e3 = op(e2, e2)) & (e0 = op(op(e2, e2), e2)) & (e1 = op(op(e2, e2), op(e2, e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG154+1.p', ax20)).
fof(f1590, plain, (~ spl66_159 | ~ spl66_161 | ~ spl66_22), inference(avatar_split_clause, [], [f590, f691, f1587, f1578])).
fof(f590, plain, (~ (e1 = op(e2, e2)) | ~ (e3 = op(op(e2, e2), e2)) | ~ (e0 = op(op(e2, e2), op(e2, e2)))), inference(cnf_transformation, [], [f43])).
fof(f43, plain, (~ (e1 = op(e2, e2)) | ~ (e3 = op(op(e2, e2), e2)) | ~ (e0 = op(op(e2, e2), op(e2, e2)))), inference(ennf_transformation, [], [f19])).
fof(f19, plain, ~ ((e1 = op(e2, e2)) & (e3 = op(op(e2, e2), e2)) & (e0 = op(op(e2, e2), op(e2, e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG154+1.p', ax19)).
fof(f1585, plain, (~ spl66_159 | ~ spl66_160 | ~ spl66_24), inference(avatar_split_clause, [], [f589, f699, f1582, f1578])).
fof(f589, plain, (~ (e3 = op(e2, e2)) | ~ (e1 = op(op(e2, e2), e2)) | ~ (e0 = op(op(e2, e2), op(e2, e2)))), inference(cnf_transformation, [], [f42])).
fof(f42, plain, (~ (e3 = op(e2, e2)) | ~ (e1 = op(op(e2, e2), e2)) | ~ (e0 = op(op(e2, e2), op(e2, e2)))), inference(ennf_transformation, [], [f18])).
fof(f18, plain, ~ ((e3 = op(e2, e2)) & (e1 = op(op(e2, e2), e2)) & (e0 = op(op(e2, e2), op(e2, e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG154+1.p', ax18)).
fof(f1576, plain, (~ spl66_158 | ~ spl66_154 | ~ spl66_41), inference(avatar_split_clause, [], [f588, f772, f1552, f1572])).
fof(f588, plain, (~ (e0 = op(e1, e1)) | ~ (e2 = op(op(e1, e1), e1)) | ~ (e3 = op(op(e1, e1), op(e1, e1)))), inference(cnf_transformation, [], [f41])).
fof(f41, plain, (~ (e0 = op(e1, e1)) | ~ (e2 = op(op(e1, e1), e1)) | ~ (e3 = op(op(e1, e1), op(e1, e1)))), inference(ennf_transformation, [], [f17])).
fof(f17, plain, ~ ((e0 = op(e1, e1)) & (e2 = op(op(e1, e1), e1)) & (e3 = op(op(e1, e1), op(e1, e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG154+1.p', ax17)).
fof(f1575, plain, (~ spl66_158 | ~ spl66_157 | ~ spl66_43), inference(avatar_split_clause, [], [f587, f780, f1566, f1572])).
fof(f587, plain, (~ (e2 = op(e1, e1)) | ~ (e0 = op(op(e1, e1), e1)) | ~ (e3 = op(op(e1, e1), op(e1, e1)))), inference(cnf_transformation, [], [f40])).
fof(f40, plain, (~ (e2 = op(e1, e1)) | ~ (e0 = op(op(e1, e1), e1)) | ~ (e3 = op(op(e1, e1), op(e1, e1)))), inference(ennf_transformation, [], [f16])).
fof(f16, plain, ~ ((e2 = op(e1, e1)) & (e0 = op(op(e1, e1), e1)) & (e3 = op(op(e1, e1), op(e1, e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG154+1.p', ax16)).
fof(f1570, plain, (~ spl66_156 | ~ spl66_155 | ~ spl66_41), inference(avatar_split_clause, [], [f586, f772, f1557, f1562])).
fof(f586, plain, (~ (e0 = op(e1, e1)) | ~ (e3 = op(op(e1, e1), e1)) | ~ (e2 = op(op(e1, e1), op(e1, e1)))), inference(cnf_transformation, [], [f39])).
fof(f39, plain, (~ (e0 = op(e1, e1)) | ~ (e3 = op(op(e1, e1), e1)) | ~ (e2 = op(op(e1, e1), op(e1, e1)))), inference(ennf_transformation, [], [f15])).
fof(f15, plain, ~ ((e0 = op(e1, e1)) & (e3 = op(op(e1, e1), e1)) & (e2 = op(op(e1, e1), op(e1, e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG154+1.p', ax15)).
fof(f1569, plain, (~ spl66_156 | ~ spl66_157 | ~ spl66_44), inference(avatar_split_clause, [], [f585, f784, f1566, f1562])).
fof(f585, plain, (~ (e3 = op(e1, e1)) | ~ (e0 = op(op(e1, e1), e1)) | ~ (e2 = op(op(e1, e1), op(e1, e1)))), inference(cnf_transformation, [], [f38])).
fof(f38, plain, (~ (e3 = op(e1, e1)) | ~ (e0 = op(op(e1, e1), e1)) | ~ (e2 = op(op(e1, e1), op(e1, e1)))), inference(ennf_transformation, [], [f14])).
fof(f14, plain, ~ ((e3 = op(e1, e1)) & (e0 = op(op(e1, e1), e1)) & (e2 = op(op(e1, e1), op(e1, e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG154+1.p', ax14)).
fof(f1560, plain, (~ spl66_153 | ~ spl66_155 | ~ spl66_43), inference(avatar_split_clause, [], [f584, f780, f1557, f1548])).
fof(f584, plain, (~ (e2 = op(e1, e1)) | ~ (e3 = op(op(e1, e1), e1)) | ~ (e0 = op(op(e1, e1), op(e1, e1)))), inference(cnf_transformation, [], [f37])).
fof(f37, plain, (~ (e2 = op(e1, e1)) | ~ (e3 = op(op(e1, e1), e1)) | ~ (e0 = op(op(e1, e1), op(e1, e1)))), inference(ennf_transformation, [], [f13])).
fof(f13, plain, ~ ((e2 = op(e1, e1)) & (e3 = op(op(e1, e1), e1)) & (e0 = op(op(e1, e1), op(e1, e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG154+1.p', ax13)).
fof(f1555, plain, (~ spl66_153 | ~ spl66_154 | ~ spl66_44), inference(avatar_split_clause, [], [f583, f784, f1552, f1548])).
fof(f583, plain, (~ (e3 = op(e1, e1)) | ~ (e2 = op(op(e1, e1), e1)) | ~ (e0 = op(op(e1, e1), op(e1, e1)))), inference(cnf_transformation, [], [f36])).
fof(f36, plain, (~ (e3 = op(e1, e1)) | ~ (e2 = op(op(e1, e1), e1)) | ~ (e0 = op(op(e1, e1), op(e1, e1)))), inference(ennf_transformation, [], [f12])).
fof(f12, plain, ~ ((e3 = op(e1, e1)) & (e2 = op(op(e1, e1), e1)) & (e0 = op(op(e1, e1), op(e1, e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG154+1.p', ax12)).
fof(f1546, plain, (~ spl66_152 | ~ spl66_148 | ~ spl66_62), inference(avatar_split_clause, [], [f582, f861, f1522, f1542])).
fof(f582, plain, (~ (op(e0, e0) = e1) | ~ (e2 = op(op(e0, e0), e0)) | ~ (e3 = op(op(e0, e0), op(e0, e0)))), inference(cnf_transformation, [], [f35])).
fof(f35, plain, (~ (op(e0, e0) = e1) | ~ (e2 = op(op(e0, e0), e0)) | ~ (e3 = op(op(e0, e0), op(e0, e0)))), inference(ennf_transformation, [], [f11])).
fof(f11, plain, ~ ((op(e0, e0) = e1) & (e2 = op(op(e0, e0), e0)) & (e3 = op(op(e0, e0), op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG154+1.p', ax11)).
fof(f1545, plain, (~ spl66_152 | ~ spl66_151 | ~ spl66_63), inference(avatar_split_clause, [], [f581, f865, f1536, f1542])).
fof(f581, plain, (~ (op(e0, e0) = e2) | ~ (e1 = op(op(e0, e0), e0)) | ~ (e3 = op(op(e0, e0), op(e0, e0)))), inference(cnf_transformation, [], [f34])).
fof(f34, plain, (~ (op(e0, e0) = e2) | ~ (e1 = op(op(e0, e0), e0)) | ~ (e3 = op(op(e0, e0), op(e0, e0)))), inference(ennf_transformation, [], [f10])).
fof(f10, plain, ~ ((op(e0, e0) = e2) & (e1 = op(op(e0, e0), e0)) & (e3 = op(op(e0, e0), op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG154+1.p', ax10)).
fof(f1540, plain, (~ spl66_150 | ~ spl66_149 | ~ spl66_62), inference(avatar_split_clause, [], [f580, f861, f1527, f1532])).
fof(f580, plain, (~ (op(e0, e0) = e1) | ~ (e3 = op(op(e0, e0), e0)) | ~ (e2 = op(op(e0, e0), op(e0, e0)))), inference(cnf_transformation, [], [f33])).
fof(f33, plain, (~ (op(e0, e0) = e1) | ~ (e3 = op(op(e0, e0), e0)) | ~ (e2 = op(op(e0, e0), op(e0, e0)))), inference(ennf_transformation, [], [f9])).
fof(f9, plain, ~ ((op(e0, e0) = e1) & (e3 = op(op(e0, e0), e0)) & (e2 = op(op(e0, e0), op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG154+1.p', ax9)).
fof(f1539, plain, (~ spl66_150 | ~ spl66_151 | ~ spl66_64), inference(avatar_split_clause, [], [f579, f869, f1536, f1532])).
fof(f579, plain, (~ (op(e0, e0) = e3) | ~ (e1 = op(op(e0, e0), e0)) | ~ (e2 = op(op(e0, e0), op(e0, e0)))), inference(cnf_transformation, [], [f32])).
fof(f32, plain, (~ (op(e0, e0) = e3) | ~ (e1 = op(op(e0, e0), e0)) | ~ (e2 = op(op(e0, e0), op(e0, e0)))), inference(ennf_transformation, [], [f8])).
fof(f8, plain, ~ ((op(e0, e0) = e3) & (e1 = op(op(e0, e0), e0)) & (e2 = op(op(e0, e0), op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG154+1.p', ax8)).
fof(f1530, plain, (~ spl66_147 | ~ spl66_149 | ~ spl66_63), inference(avatar_split_clause, [], [f578, f865, f1527, f1518])).
fof(f578, plain, (~ (op(e0, e0) = e2) | ~ (e3 = op(op(e0, e0), e0)) | ~ (e1 = op(op(e0, e0), op(e0, e0)))), inference(cnf_transformation, [], [f31])).
fof(f31, plain, (~ (op(e0, e0) = e2) | ~ (e3 = op(op(e0, e0), e0)) | ~ (e1 = op(op(e0, e0), op(e0, e0)))), inference(ennf_transformation, [], [f7])).
fof(f7, plain, ~ ((op(e0, e0) = e2) & (e3 = op(op(e0, e0), e0)) & (e1 = op(op(e0, e0), op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG154+1.p', ax7)).
fof(f1525, plain, (~ spl66_147 | ~ spl66_148 | ~ spl66_64), inference(avatar_split_clause, [], [f577, f869, f1522, f1518])).
fof(f577, plain, (~ (op(e0, e0) = e3) | ~ (e2 = op(op(e0, e0), e0)) | ~ (e1 = op(op(e0, e0), op(e0, e0)))), inference(cnf_transformation, [], [f30])).
fof(f30, plain, (~ (op(e0, e0) = e3) | ~ (e2 = op(op(e0, e0), e0)) | ~ (e1 = op(op(e0, e0), op(e0, e0)))), inference(ennf_transformation, [], [f6])).
fof(f6, plain, ~ ((op(e0, e0) = e3) & (e2 = op(op(e0, e0), e0)) & (e1 = op(op(e0, e0), op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG154+1.p', ax6)).
fof(f1516, plain, (spl66_142 | spl66_141 | spl66_140 | spl66_139 | spl66_138 | spl66_137 | spl66_136 | spl66_135 | spl66_134 | spl66_133 | spl66_132 | spl66_131 | spl66_130 | spl66_129 | spl66_128 | spl66_127 | spl66_126 | spl66_125 | spl66_124 | spl66_123 | spl66_122 | spl66_121 | spl66_120 | spl66_119 | spl66_118 | spl66_117 | spl66_116 | spl66_115 | spl66_114 | spl66_113 | spl66_112 | spl66_111 | spl66_110 | spl66_109 | spl66_108 | spl66_107 | spl66_106 | spl66_105 | spl66_104 | spl66_103 | spl66_102 | spl66_101 | spl66_100 | spl66_99 | spl66_98 | spl66_97 | spl66_96 | spl66_95 | spl66_94 | spl66_93 | spl66_92 | spl66_91 | spl66_90 | spl66_89 | spl66_88 | spl66_87 | spl66_86 | spl66_85 | spl66_84 | spl66_83 | spl66_82 | spl66_81 | spl66_80 | spl66_4), inference(avatar_split_clause, [], [f565, f614, f987, f995, f1003, f1011, f1019, f1027, f1035, f1043, f1051, f1059, f1067, f1075, f1083, f1091, f1099, f1107, f1115, f1123, f1131, f1139, f1147, f1155, f1163, f1171, f1179, f1187, f1195, f1203, f1211, f1219, f1227, f1235, f1243, f1251, f1259, f1267, f1275, f1283, f1291, f1299, f1307, f1315, f1323, f1331, f1339, f1347, f1355, f1363, f1371, f1379, f1387, f1395, f1403, f1411, f1419, f1427, f1435, f1443, f1451, f1459, f1467, f1475, f1483])).
fof(f1483, plain, (spl66_142 <=> sP0), introduced(avatar_definition, [new_symbols(naming, [spl66_142])])).
fof(f1475, plain, (spl66_141 <=> sP1), introduced(avatar_definition, [new_symbols(naming, [spl66_141])])).
fof(f1467, plain, (spl66_140 <=> sP2), introduced(avatar_definition, [new_symbols(naming, [spl66_140])])).
fof(f1459, plain, (spl66_139 <=> sP3), introduced(avatar_definition, [new_symbols(naming, [spl66_139])])).
fof(f1451, plain, (spl66_138 <=> sP4), introduced(avatar_definition, [new_symbols(naming, [spl66_138])])).
fof(f1443, plain, (spl66_137 <=> sP5), introduced(avatar_definition, [new_symbols(naming, [spl66_137])])).
fof(f1435, plain, (spl66_136 <=> sP6), introduced(avatar_definition, [new_symbols(naming, [spl66_136])])).
fof(f1427, plain, (spl66_135 <=> sP7), introduced(avatar_definition, [new_symbols(naming, [spl66_135])])).
fof(f1419, plain, (spl66_134 <=> sP8), introduced(avatar_definition, [new_symbols(naming, [spl66_134])])).
fof(f1411, plain, (spl66_133 <=> sP9), introduced(avatar_definition, [new_symbols(naming, [spl66_133])])).
fof(f1403, plain, (spl66_132 <=> sP10), introduced(avatar_definition, [new_symbols(naming, [spl66_132])])).
fof(f1395, plain, (spl66_131 <=> sP11), introduced(avatar_definition, [new_symbols(naming, [spl66_131])])).
fof(f1387, plain, (spl66_130 <=> sP12), introduced(avatar_definition, [new_symbols(naming, [spl66_130])])).
fof(f1379, plain, (spl66_129 <=> sP13), introduced(avatar_definition, [new_symbols(naming, [spl66_129])])).
fof(f1371, plain, (spl66_128 <=> sP14), introduced(avatar_definition, [new_symbols(naming, [spl66_128])])).
fof(f1363, plain, (spl66_127 <=> sP15), introduced(avatar_definition, [new_symbols(naming, [spl66_127])])).
fof(f1355, plain, (spl66_126 <=> sP16), introduced(avatar_definition, [new_symbols(naming, [spl66_126])])).
fof(f1347, plain, (spl66_125 <=> sP17), introduced(avatar_definition, [new_symbols(naming, [spl66_125])])).
fof(f1339, plain, (spl66_124 <=> sP18), introduced(avatar_definition, [new_symbols(naming, [spl66_124])])).
fof(f1331, plain, (spl66_123 <=> sP19), introduced(avatar_definition, [new_symbols(naming, [spl66_123])])).
fof(f1323, plain, (spl66_122 <=> sP20), introduced(avatar_definition, [new_symbols(naming, [spl66_122])])).
fof(f1315, plain, (spl66_121 <=> sP21), introduced(avatar_definition, [new_symbols(naming, [spl66_121])])).
fof(f1307, plain, (spl66_120 <=> sP22), introduced(avatar_definition, [new_symbols(naming, [spl66_120])])).
fof(f1299, plain, (spl66_119 <=> sP23), introduced(avatar_definition, [new_symbols(naming, [spl66_119])])).
fof(f1291, plain, (spl66_118 <=> sP24), introduced(avatar_definition, [new_symbols(naming, [spl66_118])])).
fof(f1283, plain, (spl66_117 <=> sP25), introduced(avatar_definition, [new_symbols(naming, [spl66_117])])).
fof(f1275, plain, (spl66_116 <=> sP26), introduced(avatar_definition, [new_symbols(naming, [spl66_116])])).
fof(f1267, plain, (spl66_115 <=> sP27), introduced(avatar_definition, [new_symbols(naming, [spl66_115])])).
fof(f1259, plain, (spl66_114 <=> sP28), introduced(avatar_definition, [new_symbols(naming, [spl66_114])])).
fof(f1251, plain, (spl66_113 <=> sP29), introduced(avatar_definition, [new_symbols(naming, [spl66_113])])).
fof(f1243, plain, (spl66_112 <=> sP30), introduced(avatar_definition, [new_symbols(naming, [spl66_112])])).
fof(f1235, plain, (spl66_111 <=> sP31), introduced(avatar_definition, [new_symbols(naming, [spl66_111])])).
fof(f1227, plain, (spl66_110 <=> sP32), introduced(avatar_definition, [new_symbols(naming, [spl66_110])])).
fof(f1219, plain, (spl66_109 <=> sP33), introduced(avatar_definition, [new_symbols(naming, [spl66_109])])).
fof(f1211, plain, (spl66_108 <=> sP34), introduced(avatar_definition, [new_symbols(naming, [spl66_108])])).
fof(f1203, plain, (spl66_107 <=> sP35), introduced(avatar_definition, [new_symbols(naming, [spl66_107])])).
fof(f1195, plain, (spl66_106 <=> sP36), introduced(avatar_definition, [new_symbols(naming, [spl66_106])])).
fof(f1187, plain, (spl66_105 <=> sP37), introduced(avatar_definition, [new_symbols(naming, [spl66_105])])).
fof(f1179, plain, (spl66_104 <=> sP38), introduced(avatar_definition, [new_symbols(naming, [spl66_104])])).
fof(f1171, plain, (spl66_103 <=> sP39), introduced(avatar_definition, [new_symbols(naming, [spl66_103])])).
fof(f1163, plain, (spl66_102 <=> sP40), introduced(avatar_definition, [new_symbols(naming, [spl66_102])])).
fof(f1155, plain, (spl66_101 <=> sP41), introduced(avatar_definition, [new_symbols(naming, [spl66_101])])).
fof(f1147, plain, (spl66_100 <=> sP42), introduced(avatar_definition, [new_symbols(naming, [spl66_100])])).
fof(f1139, plain, (spl66_99 <=> sP43), introduced(avatar_definition, [new_symbols(naming, [spl66_99])])).
fof(f1131, plain, (spl66_98 <=> sP44), introduced(avatar_definition, [new_symbols(naming, [spl66_98])])).
fof(f1123, plain, (spl66_97 <=> sP45), introduced(avatar_definition, [new_symbols(naming, [spl66_97])])).
fof(f1115, plain, (spl66_96 <=> sP46), introduced(avatar_definition, [new_symbols(naming, [spl66_96])])).
fof(f1107, plain, (spl66_95 <=> sP47), introduced(avatar_definition, [new_symbols(naming, [spl66_95])])).
fof(f1099, plain, (spl66_94 <=> sP48), introduced(avatar_definition, [new_symbols(naming, [spl66_94])])).
fof(f1091, plain, (spl66_93 <=> sP49), introduced(avatar_definition, [new_symbols(naming, [spl66_93])])).
fof(f1083, plain, (spl66_92 <=> sP50), introduced(avatar_definition, [new_symbols(naming, [spl66_92])])).
fof(f1075, plain, (spl66_91 <=> sP51), introduced(avatar_definition, [new_symbols(naming, [spl66_91])])).
fof(f1067, plain, (spl66_90 <=> sP52), introduced(avatar_definition, [new_symbols(naming, [spl66_90])])).
fof(f1059, plain, (spl66_89 <=> sP53), introduced(avatar_definition, [new_symbols(naming, [spl66_89])])).
fof(f1051, plain, (spl66_88 <=> sP54), introduced(avatar_definition, [new_symbols(naming, [spl66_88])])).
fof(f1043, plain, (spl66_87 <=> sP55), introduced(avatar_definition, [new_symbols(naming, [spl66_87])])).
fof(f1035, plain, (spl66_86 <=> sP56), introduced(avatar_definition, [new_symbols(naming, [spl66_86])])).
fof(f1027, plain, (spl66_85 <=> sP57), introduced(avatar_definition, [new_symbols(naming, [spl66_85])])).
fof(f1019, plain, (spl66_84 <=> sP58), introduced(avatar_definition, [new_symbols(naming, [spl66_84])])).
fof(f1011, plain, (spl66_83 <=> sP59), introduced(avatar_definition, [new_symbols(naming, [spl66_83])])).
fof(f1003, plain, (spl66_82 <=> sP60), introduced(avatar_definition, [new_symbols(naming, [spl66_82])])).
fof(f995, plain, (spl66_81 <=> sP61), introduced(avatar_definition, [new_symbols(naming, [spl66_81])])).
fof(f987, plain, (spl66_80 <=> sP62), introduced(avatar_definition, [new_symbols(naming, [spl66_80])])).
fof(f565, plain, ((e3 = op(e3, e3)) | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f120])).
fof(f120, plain, (((((e3 = op(e3, e3)) | ~ (e3 = op(e3, e3))) & ((e2 = op(e2, e3)) | ~ (e3 = op(e2, e2))) & ((e1 = op(e1, e3)) | ~ (e3 = op(e1, e1))) & ((e0 = op(e0, e3)) | ~ (op(e0, e0) = e3)) & (e3 = op(e3, op(e3, e3))) & (e2 = op(e3, op(e3, e2))) & (e1 = op(e3, op(e3, e1))) & (e0 = op(e3, op(e3, e0)))) | sP65 | sP64 | sP63) & ((~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0)), inference(definition_folding, [], [f5, e119, e118, e117, e116, e115, e114, e113, e112, e111, e110, e109, e108, e107, e106, e105, e104, e103, e102, e101, e100, e99, e98, e97, e96, e95, e94, e93, e92, e91, e90, e89, e88, e87, e86, e85, e84, e83, e82, e81, e80, e79, e78, e77, e76, e75, e74, e73, e72, e71, e70, e69, e68, e67, e66, e65, e64, e63, e62, e61, e60, e59, e58, e57, e56, e55, e54])).
fof(f54, plain, ((~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP0), inference(usedef, [], [e54])).
fof(e54, plain, (sP0 <=> (~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f55, plain, ((~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP1), inference(usedef, [], [e55])).
fof(e55, plain, (sP1 <=> (~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f56, plain, ((~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP2), inference(usedef, [], [e56])).
fof(e56, plain, (sP2 <=> (~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f57, plain, ((~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP3), inference(usedef, [], [e57])).
fof(e57, plain, (sP3 <=> (~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f58, plain, ((~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | ~ sP4), inference(usedef, [], [e58])).
fof(e58, plain, (sP4 <=> (~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f59, plain, ((~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | ~ sP5), inference(usedef, [], [e59])).
fof(e59, plain, (sP5 <=> (~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f60, plain, ((~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | ~ sP6), inference(usedef, [], [e60])).
fof(e60, plain, (sP6 <=> (~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f61, plain, ((~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | ~ sP7), inference(usedef, [], [e61])).
fof(e61, plain, (sP7 <=> (~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f62, plain, ((~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | ~ sP8), inference(usedef, [], [e62])).
fof(e62, plain, (sP8 <=> (~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f63, plain, ((~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | ~ sP9), inference(usedef, [], [e63])).
fof(e63, plain, (sP9 <=> (~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f64, plain, ((~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | ~ sP10), inference(usedef, [], [e64])).
fof(e64, plain, (sP10 <=> (~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f65, plain, ((~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | ~ sP11), inference(usedef, [], [e65])).
fof(e65, plain, (sP11 <=> (~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f66, plain, ((~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | ~ sP12), inference(usedef, [], [e66])).
fof(e66, plain, (sP12 <=> (~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f67, plain, ((~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | ~ sP13), inference(usedef, [], [e67])).
fof(e67, plain, (sP13 <=> (~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f68, plain, ((~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | ~ sP14), inference(usedef, [], [e68])).
fof(e68, plain, (sP14 <=> (~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f69, plain, ((~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | ~ sP15), inference(usedef, [], [e69])).
fof(e69, plain, (sP15 <=> (~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f70, plain, ((~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | ~ sP16), inference(usedef, [], [e70])).
fof(e70, plain, (sP16 <=> (~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f71, plain, ((~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | ~ sP17), inference(usedef, [], [e71])).
fof(e71, plain, (sP17 <=> (~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f72, plain, ((~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | ~ sP18), inference(usedef, [], [e72])).
fof(e72, plain, (sP18 <=> (~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f73, plain, ((~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | ~ sP19), inference(usedef, [], [e73])).
fof(e73, plain, (sP19 <=> (~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP19])])).
fof(f74, plain, ((~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP20), inference(usedef, [], [e74])).
fof(e74, plain, (sP20 <=> (~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP20])])).
fof(f75, plain, ((~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP21), inference(usedef, [], [e75])).
fof(e75, plain, (sP21 <=> (~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP21])])).
fof(f76, plain, ((~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP22), inference(usedef, [], [e76])).
fof(e76, plain, (sP22 <=> (~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP22])])).
fof(f77, plain, ((~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP23), inference(usedef, [], [e77])).
fof(e77, plain, (sP23 <=> (~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP23])])).
fof(f78, plain, ((~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | ~ sP24), inference(usedef, [], [e78])).
fof(e78, plain, (sP24 <=> (~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP24])])).
fof(f79, plain, ((~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | ~ sP25), inference(usedef, [], [e79])).
fof(e79, plain, (sP25 <=> (~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP25])])).
fof(f80, plain, ((~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | ~ sP26), inference(usedef, [], [e80])).
fof(e80, plain, (sP26 <=> (~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP26])])).
fof(f81, plain, ((~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | ~ sP27), inference(usedef, [], [e81])).
fof(e81, plain, (sP27 <=> (~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP27])])).
fof(f82, plain, ((~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | ~ sP28), inference(usedef, [], [e82])).
fof(e82, plain, (sP28 <=> (~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP28])])).
fof(f83, plain, ((~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | ~ sP29), inference(usedef, [], [e83])).
fof(e83, plain, (sP29 <=> (~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP29])])).
fof(f84, plain, ((~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | ~ sP30), inference(usedef, [], [e84])).
fof(e84, plain, (sP30 <=> (~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP30])])).
fof(f85, plain, ((~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | ~ sP31), inference(usedef, [], [e85])).
fof(e85, plain, (sP31 <=> (~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP31])])).
fof(f86, plain, ((~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | ~ sP32), inference(usedef, [], [e86])).
fof(e86, plain, (sP32 <=> (~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP32])])).
fof(f87, plain, ((~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | ~ sP33), inference(usedef, [], [e87])).
fof(e87, plain, (sP33 <=> (~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP33])])).
fof(f88, plain, ((~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | ~ sP34), inference(usedef, [], [e88])).
fof(e88, plain, (sP34 <=> (~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP34])])).
fof(f89, plain, ((~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | ~ sP35), inference(usedef, [], [e89])).
fof(e89, plain, (sP35 <=> (~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP35])])).
fof(f90, plain, ((~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | ~ sP36), inference(usedef, [], [e90])).
fof(e90, plain, (sP36 <=> (~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP36])])).
fof(f91, plain, ((~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | ~ sP37), inference(usedef, [], [e91])).
fof(e91, plain, (sP37 <=> (~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP37])])).
fof(f92, plain, ((~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | ~ sP38), inference(usedef, [], [e92])).
fof(e92, plain, (sP38 <=> (~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP38])])).
fof(f93, plain, ((~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | ~ sP39), inference(usedef, [], [e93])).
fof(e93, plain, (sP39 <=> (~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP39])])).
fof(f94, plain, ((~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP40), inference(usedef, [], [e94])).
fof(e94, plain, (sP40 <=> (~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP40])])).
fof(f95, plain, ((~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP41), inference(usedef, [], [e95])).
fof(e95, plain, (sP41 <=> (~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP41])])).
fof(f96, plain, ((~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP42), inference(usedef, [], [e96])).
fof(e96, plain, (sP42 <=> (~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP42])])).
fof(f97, plain, ((~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP43), inference(usedef, [], [e97])).
fof(e97, plain, (sP43 <=> (~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP43])])).
fof(f98, plain, ((~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | ~ sP44), inference(usedef, [], [e98])).
fof(e98, plain, (sP44 <=> (~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP44])])).
fof(f99, plain, ((~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | ~ sP45), inference(usedef, [], [e99])).
fof(e99, plain, (sP45 <=> (~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP45])])).
fof(f100, plain, ((~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | ~ sP46), inference(usedef, [], [e100])).
fof(e100, plain, (sP46 <=> (~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP46])])).
fof(f101, plain, ((~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | ~ sP47), inference(usedef, [], [e101])).
fof(e101, plain, (sP47 <=> (~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP47])])).
fof(f102, plain, ((~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | ~ sP48), inference(usedef, [], [e102])).
fof(e102, plain, (sP48 <=> (~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP48])])).
fof(f103, plain, ((~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | ~ sP49), inference(usedef, [], [e103])).
fof(e103, plain, (sP49 <=> (~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP49])])).
fof(f104, plain, ((~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | ~ sP50), inference(usedef, [], [e104])).
fof(e104, plain, (sP50 <=> (~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP50])])).
fof(f105, plain, ((~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | ~ sP51), inference(usedef, [], [e105])).
fof(e105, plain, (sP51 <=> (~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP51])])).
fof(f106, plain, ((~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | ~ sP52), inference(usedef, [], [e106])).
fof(e106, plain, (sP52 <=> (~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP52])])).
fof(f107, plain, ((~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | ~ sP53), inference(usedef, [], [e107])).
fof(e107, plain, (sP53 <=> (~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP53])])).
fof(f108, plain, ((~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | ~ sP54), inference(usedef, [], [e108])).
fof(e108, plain, (sP54 <=> (~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP54])])).
fof(f109, plain, ((~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | ~ sP55), inference(usedef, [], [e109])).
fof(e109, plain, (sP55 <=> (~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP55])])).
fof(f110, plain, ((~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | ~ sP56), inference(usedef, [], [e110])).
fof(e110, plain, (sP56 <=> (~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP56])])).
fof(f111, plain, ((~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | ~ sP57), inference(usedef, [], [e111])).
fof(e111, plain, (sP57 <=> (~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP57])])).
fof(f112, plain, ((~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | ~ sP58), inference(usedef, [], [e112])).
fof(e112, plain, (sP58 <=> (~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP58])])).
fof(f113, plain, ((~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | ~ sP59), inference(usedef, [], [e113])).
fof(e113, plain, (sP59 <=> (~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP59])])).
fof(f114, plain, ((~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ~ sP60), inference(usedef, [], [e114])).
fof(e114, plain, (sP60 <=> (~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP60])])).
fof(f115, plain, ((~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ~ sP61), inference(usedef, [], [e115])).
fof(e115, plain, (sP61 <=> (~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP61])])).
fof(f116, plain, ((~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ~ sP62), inference(usedef, [], [e116])).
fof(e116, plain, (sP62 <=> (~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP62])])).
fof(f117, plain, ((((e3 = op(e3, e0)) | ~ (e0 = op(e3, e3))) & ((e2 = op(e2, e0)) | ~ (e0 = op(e2, e2))) & ((e1 = op(e1, e0)) | ~ (e0 = op(e1, e1))) & ((e0 = op(e0, e0)) | ~ (e0 = op(e0, e0))) & (e3 = op(e0, op(e0, e3))) & (e2 = op(e0, op(e0, e2))) & (e1 = op(e0, op(e0, e1))) & (e0 = op(e0, op(e0, e0)))) | ~ sP63), inference(usedef, [], [e117])).
fof(e117, plain, (sP63 <=> (((e3 = op(e3, e0)) | ~ (e0 = op(e3, e3))) & ((e2 = op(e2, e0)) | ~ (e0 = op(e2, e2))) & ((e1 = op(e1, e0)) | ~ (e0 = op(e1, e1))) & ((e0 = op(e0, e0)) | ~ (e0 = op(e0, e0))) & (e3 = op(e0, op(e0, e3))) & (e2 = op(e0, op(e0, e2))) & (e1 = op(e0, op(e0, e1))) & (e0 = op(e0, op(e0, e0))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP63])])).
fof(f118, plain, ((((e3 = op(e3, e1)) | ~ (e1 = op(e3, e3))) & ((e2 = op(e2, e1)) | ~ (e1 = op(e2, e2))) & ((e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & ((e0 = op(e0, e1)) | ~ (op(e0, e0) = e1)) & (e3 = op(e1, op(e1, e3))) & (e2 = op(e1, op(e1, e2))) & (e1 = op(e1, op(e1, e1))) & (e0 = op(e1, op(e1, e0)))) | ~ sP64), inference(usedef, [], [e118])).
fof(e118, plain, (sP64 <=> (((e3 = op(e3, e1)) | ~ (e1 = op(e3, e3))) & ((e2 = op(e2, e1)) | ~ (e1 = op(e2, e2))) & ((e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & ((e0 = op(e0, e1)) | ~ (op(e0, e0) = e1)) & (e3 = op(e1, op(e1, e3))) & (e2 = op(e1, op(e1, e2))) & (e1 = op(e1, op(e1, e1))) & (e0 = op(e1, op(e1, e0))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP64])])).
fof(f119, plain, ((((e3 = op(e3, e2)) | ~ (e2 = op(e3, e3))) & ((e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & ((e1 = op(e1, e2)) | ~ (e2 = op(e1, e1))) & ((e0 = op(e0, e2)) | ~ (op(e0, e0) = e2)) & (e3 = op(e2, op(e2, e3))) & (e2 = op(e2, op(e2, e2))) & (e1 = op(e2, op(e2, e1))) & (e0 = op(e2, op(e2, e0)))) | ~ sP65), inference(usedef, [], [e119])).
fof(e119, plain, (sP65 <=> (((e3 = op(e3, e2)) | ~ (e2 = op(e3, e3))) & ((e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & ((e1 = op(e1, e2)) | ~ (e2 = op(e1, e1))) & ((e0 = op(e0, e2)) | ~ (op(e0, e0) = e2)) & (e3 = op(e2, op(e2, e3))) & (e2 = op(e2, op(e2, e2))) & (e1 = op(e2, op(e2, e1))) & (e0 = op(e2, op(e2, e0))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP65])])).
fof(f5, plain, (((((e3 = op(e3, e3)) | ~ (e3 = op(e3, e3))) & ((e2 = op(e2, e3)) | ~ (e3 = op(e2, e2))) & ((e1 = op(e1, e3)) | ~ (e3 = op(e1, e1))) & ((e0 = op(e0, e3)) | ~ (op(e0, e0) = e3)) & (e3 = op(e3, op(e3, e3))) & (e2 = op(e3, op(e3, e2))) & (e1 = op(e3, op(e3, e1))) & (e0 = op(e3, op(e3, e0)))) | (((e3 = op(e3, e2)) | ~ (e2 = op(e3, e3))) & ((e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & ((e1 = op(e1, e2)) | ~ (e2 = op(e1, e1))) & ((e0 = op(e0, e2)) | ~ (op(e0, e0) = e2)) & (e3 = op(e2, op(e2, e3))) & (e2 = op(e2, op(e2, e2))) & (e1 = op(e2, op(e2, e1))) & (e0 = op(e2, op(e2, e0)))) | (((e3 = op(e3, e1)) | ~ (e1 = op(e3, e3))) & ((e2 = op(e2, e1)) | ~ (e1 = op(e2, e2))) & ((e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & ((e0 = op(e0, e1)) | ~ (op(e0, e0) = e1)) & (e3 = op(e1, op(e1, e3))) & (e2 = op(e1, op(e1, e2))) & (e1 = op(e1, op(e1, e1))) & (e0 = op(e1, op(e1, e0)))) | (((e3 = op(e3, e0)) | ~ (e0 = op(e3, e3))) & ((e2 = op(e2, e0)) | ~ (e0 = op(e2, e2))) & ((e1 = op(e1, e0)) | ~ (e0 = op(e1, e1))) & ((e0 = op(e0, e0)) | ~ (e0 = op(e0, e0))) & (e3 = op(e0, op(e0, e3))) & (e2 = op(e0, op(e0, e2))) & (e1 = op(e0, op(e0, e1))) & (e0 = op(e0, op(e0, e0))))) & ((~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | (~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | (~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | (~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | (~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | (~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | (~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | (~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | (~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | (~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | (~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | (~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | (~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | (~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | (~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | (~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | (~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | (~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | (~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | (~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | (~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | (~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | (~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | (~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | (~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | (~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | (~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | (~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | (~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | (~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | (~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | (~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | (~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | (~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | (~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | (~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | (~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | (~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | (~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | (~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | (~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | (~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | (~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | (~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | (~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | (~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | (~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | (~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | (~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | (~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | (~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | (~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | (~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | (~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | (~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | (~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | (~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | (~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | (~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | (~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | (~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | (~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | (~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | (~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG154+1.p', ax5)).
fof(f1515, plain, (spl66_142 | spl66_141 | spl66_140 | spl66_139 | spl66_138 | spl66_137 | spl66_136 | spl66_135 | spl66_134 | spl66_133 | spl66_132 | spl66_131 | spl66_130 | spl66_129 | spl66_128 | spl66_127 | spl66_126 | spl66_125 | spl66_124 | spl66_123 | spl66_122 | spl66_121 | spl66_120 | spl66_119 | spl66_118 | spl66_117 | spl66_116 | spl66_115 | spl66_114 | spl66_113 | spl66_112 | spl66_111 | spl66_110 | spl66_109 | spl66_108 | spl66_107 | spl66_106 | spl66_105 | spl66_104 | spl66_103 | spl66_102 | spl66_101 | spl66_100 | spl66_99 | spl66_98 | spl66_97 | spl66_96 | spl66_95 | spl66_94 | spl66_93 | spl66_92 | spl66_91 | spl66_90 | spl66_89 | spl66_88 | spl66_87 | spl66_86 | spl66_85 | spl66_84 | spl66_83 | spl66_82 | spl66_81 | spl66_80 | ~ spl66_4), inference(avatar_split_clause, [], [f566, f614, f987, f995, f1003, f1011, f1019, f1027, f1035, f1043, f1051, f1059, f1067, f1075, f1083, f1091, f1099, f1107, f1115, f1123, f1131, f1139, f1147, f1155, f1163, f1171, f1179, f1187, f1195, f1203, f1211, f1219, f1227, f1235, f1243, f1251, f1259, f1267, f1275, f1283, f1291, f1299, f1307, f1315, f1323, f1331, f1339, f1347, f1355, f1363, f1371, f1379, f1387, f1395, f1403, f1411, f1419, f1427, f1435, f1443, f1451, f1459, f1467, f1475, f1483])).
fof(f566, plain, (~ (e3 = op(e3, e3)) | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f120])).
fof(f1512, plain, (spl66_75 | spl66_70 | spl66_65 | spl66_146), inference(avatar_split_clause, [], [f569, f1509, f906, f933, f960])).
fof(f960, plain, (spl66_75 <=> sP63), introduced(avatar_definition, [new_symbols(naming, [spl66_75])])).
fof(f933, plain, (spl66_70 <=> sP64), introduced(avatar_definition, [new_symbols(naming, [spl66_70])])).
fof(f906, plain, (spl66_65 <=> sP65), introduced(avatar_definition, [new_symbols(naming, [spl66_65])])).
fof(f569, plain, ((e0 = op(e3, op(e3, e0))) | sP65 | sP64 | sP63), inference(cnf_transformation, [], [f120])).
fof(f1507, plain, (spl66_75 | spl66_70 | spl66_65 | spl66_145), inference(avatar_split_clause, [], [f570, f1504, f906, f933, f960])).
fof(f570, plain, ((e1 = op(e3, op(e3, e1))) | sP65 | sP64 | sP63), inference(cnf_transformation, [], [f120])).
fof(f1502, plain, (spl66_75 | spl66_70 | spl66_65 | spl66_144), inference(avatar_split_clause, [], [f571, f1499, f906, f933, f960])).
fof(f571, plain, ((e2 = op(e3, op(e3, e2))) | sP65 | sP64 | sP63), inference(cnf_transformation, [], [f120])).
fof(f1497, plain, (spl66_75 | spl66_70 | spl66_65 | spl66_143), inference(avatar_split_clause, [], [f572, f1494, f906, f933, f960])).
fof(f572, plain, ((e3 = op(e3, op(e3, e3))) | sP65 | sP64 | sP63), inference(cnf_transformation, [], [f120])).
fof(f1492, plain, (spl66_75 | spl66_70 | spl66_65 | ~ spl66_64 | spl66_49), inference(avatar_split_clause, [], [f573, f806, f869, f906, f933, f960])).
fof(f573, plain, ((e0 = op(e0, e3)) | ~ (op(e0, e0) = e3) | sP65 | sP64 | sP63), inference(cnf_transformation, [], [f120])).
fof(f1491, plain, (spl66_75 | spl66_70 | spl66_65 | ~ spl66_44 | spl66_34), inference(avatar_split_clause, [], [f574, f742, f784, f906, f933, f960])).
fof(f574, plain, ((e1 = op(e1, e3)) | ~ (e3 = op(e1, e1)) | sP65 | sP64 | sP63), inference(cnf_transformation, [], [f120])).
fof(f1490, plain, (spl66_75 | spl66_70 | spl66_65 | ~ spl66_24 | spl66_19), inference(avatar_split_clause, [], [f575, f678, f699, f906, f933, f960])).
fof(f575, plain, ((e2 = op(e2, e3)) | ~ (e3 = op(e2, e2)) | sP65 | sP64 | sP63), inference(cnf_transformation, [], [f120])).
fof(f1489, plain, (~ spl66_142 | spl66_61), inference(avatar_split_clause, [], [f561, f857, f1483])).
fof(f561, plain, ((e0 = op(e0, e0)) | ~ sP0), inference(cnf_transformation, [], [f186])).
fof(f186, plain, ((~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP0), inference(nnf_transformation, [], [f54])).
fof(f1488, plain, (~ spl66_142 | ~ spl66_61), inference(avatar_split_clause, [], [f562, f857, f1483])).
fof(f562, plain, (~ (e0 = op(e0, e0)) | ~ sP0), inference(cnf_transformation, [], [f186])).
fof(f1481, plain, (~ spl66_141 | spl66_61), inference(avatar_split_clause, [], [f557, f857, f1475])).
fof(f557, plain, ((e0 = op(e0, e0)) | ~ sP1), inference(cnf_transformation, [], [f185])).
fof(f185, plain, ((~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP1), inference(nnf_transformation, [], [f55])).
fof(f1480, plain, (~ spl66_141 | ~ spl66_61), inference(avatar_split_clause, [], [f558, f857, f1475])).
fof(f558, plain, (~ (e0 = op(e0, e0)) | ~ sP1), inference(cnf_transformation, [], [f185])).
fof(f1473, plain, (~ spl66_140 | spl66_61), inference(avatar_split_clause, [], [f553, f857, f1467])).
fof(f553, plain, ((e0 = op(e0, e0)) | ~ sP2), inference(cnf_transformation, [], [f184])).
fof(f184, plain, ((~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP2), inference(nnf_transformation, [], [f56])).
fof(f1472, plain, (~ spl66_140 | ~ spl66_61), inference(avatar_split_clause, [], [f554, f857, f1467])).
fof(f554, plain, (~ (e0 = op(e0, e0)) | ~ sP2), inference(cnf_transformation, [], [f184])).
fof(f1465, plain, (~ spl66_139 | spl66_61), inference(avatar_split_clause, [], [f549, f857, f1459])).
fof(f549, plain, ((e0 = op(e0, e0)) | ~ sP3), inference(cnf_transformation, [], [f183])).
fof(f183, plain, ((~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP3), inference(nnf_transformation, [], [f57])).
fof(f1464, plain, (~ spl66_139 | ~ spl66_61), inference(avatar_split_clause, [], [f550, f857, f1459])).
fof(f550, plain, (~ (e0 = op(e0, e0)) | ~ sP3), inference(cnf_transformation, [], [f183])).
fof(f1456, plain, (~ spl66_138 | ~ spl66_62), inference(avatar_split_clause, [], [f546, f861, f1451])).
fof(f546, plain, (~ (op(e0, e0) = e1) | ~ sP4), inference(cnf_transformation, [], [f182])).
fof(f182, plain, ((~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | ~ sP4), inference(nnf_transformation, [], [f58])).
fof(f1455, plain, (~ spl66_138 | spl66_62), inference(avatar_split_clause, [], [f547, f861, f1451])).
fof(f547, plain, ((op(e0, e0) = e1) | ~ sP4), inference(cnf_transformation, [], [f182])).
fof(f1447, plain, (~ spl66_137 | spl66_42), inference(avatar_split_clause, [], [f543, f776, f1443])).
fof(f543, plain, ((e1 = op(e1, e1)) | ~ sP5), inference(cnf_transformation, [], [f181])).
fof(f181, plain, ((~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | ~ sP5), inference(nnf_transformation, [], [f59])).
fof(f1446, plain, (~ spl66_137 | ~ spl66_42), inference(avatar_split_clause, [], [f544, f776, f1443])).
fof(f544, plain, (~ (e1 = op(e1, e1)) | ~ sP5), inference(cnf_transformation, [], [f181])).
fof(f1441, plain, (~ spl66_136 | spl66_57), inference(avatar_split_clause, [], [f537, f840, f1435])).
fof(f537, plain, ((e0 = op(e0, e1)) | ~ sP6), inference(cnf_transformation, [], [f180])).
fof(f180, plain, ((~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | ~ sP6), inference(nnf_transformation, [], [f60])).
fof(f1440, plain, (~ spl66_136 | ~ spl66_62), inference(avatar_split_clause, [], [f538, f861, f1435])).
fof(f538, plain, (~ (op(e0, e0) = e1) | ~ sP6), inference(cnf_transformation, [], [f180])).
fof(f1439, plain, (~ spl66_136 | spl66_22), inference(avatar_split_clause, [], [f539, f691, f1435])).
fof(f539, plain, ((e1 = op(e2, e2)) | ~ sP6), inference(cnf_transformation, [], [f180])).
fof(f1438, plain, (~ spl66_136 | ~ spl66_27), inference(avatar_split_clause, [], [f540, f712, f1435])).
fof(f540, plain, (~ (e2 = op(e2, e1)) | ~ sP6), inference(cnf_transformation, [], [f180])).
fof(f1433, plain, (~ spl66_135 | spl66_57), inference(avatar_split_clause, [], [f533, f840, f1427])).
fof(f533, plain, ((e0 = op(e0, e1)) | ~ sP7), inference(cnf_transformation, [], [f179])).
fof(f179, plain, ((~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | ~ sP7), inference(nnf_transformation, [], [f61])).
fof(f1432, plain, (~ spl66_135 | ~ spl66_62), inference(avatar_split_clause, [], [f534, f861, f1427])).
fof(f534, plain, (~ (op(e0, e0) = e1) | ~ sP7), inference(cnf_transformation, [], [f179])).
fof(f1431, plain, (~ spl66_135 | spl66_2), inference(avatar_split_clause, [], [f535, f606, f1427])).
fof(f535, plain, ((e1 = op(e3, e3)) | ~ sP7), inference(cnf_transformation, [], [f179])).
fof(f1430, plain, (~ spl66_135 | ~ spl66_12), inference(avatar_split_clause, [], [f536, f648, f1427])).
fof(f536, plain, (~ (e3 = op(e3, e1)) | ~ sP7), inference(cnf_transformation, [], [f179])).
fof(f1424, plain, (~ spl66_134 | ~ spl66_63), inference(avatar_split_clause, [], [f530, f865, f1419])).
fof(f530, plain, (~ (op(e0, e0) = e2) | ~ sP8), inference(cnf_transformation, [], [f178])).
fof(f178, plain, ((~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | ~ sP8), inference(nnf_transformation, [], [f62])).
fof(f1423, plain, (~ spl66_134 | spl66_63), inference(avatar_split_clause, [], [f531, f865, f1419])).
fof(f531, plain, ((op(e0, e0) = e2) | ~ sP8), inference(cnf_transformation, [], [f178])).
fof(f1417, plain, (~ spl66_133 | spl66_53), inference(avatar_split_clause, [], [f525, f823, f1411])).
fof(f525, plain, ((e0 = op(e0, e2)) | ~ sP9), inference(cnf_transformation, [], [f177])).
fof(f177, plain, ((~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | ~ sP9), inference(nnf_transformation, [], [f63])).
fof(f1416, plain, (~ spl66_133 | ~ spl66_63), inference(avatar_split_clause, [], [f526, f865, f1411])).
fof(f526, plain, (~ (op(e0, e0) = e2) | ~ sP9), inference(cnf_transformation, [], [f177])).
fof(f1415, plain, (~ spl66_133 | spl66_43), inference(avatar_split_clause, [], [f527, f780, f1411])).
fof(f527, plain, ((e2 = op(e1, e1)) | ~ sP9), inference(cnf_transformation, [], [f177])).
fof(f1414, plain, (~ spl66_133 | ~ spl66_38), inference(avatar_split_clause, [], [f528, f759, f1411])).
fof(f528, plain, (~ (e1 = op(e1, e2)) | ~ sP9), inference(cnf_transformation, [], [f177])).
fof(f1409, plain, (~ spl66_132 | spl66_53), inference(avatar_split_clause, [], [f521, f823, f1403])).
fof(f521, plain, ((e0 = op(e0, e2)) | ~ sP10), inference(cnf_transformation, [], [f176])).
fof(f176, plain, ((~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | ~ sP10), inference(nnf_transformation, [], [f64])).
fof(f1407, plain, (~ spl66_132 | spl66_23), inference(avatar_split_clause, [], [f523, f695, f1403])).
fof(f523, plain, ((e2 = op(e2, e2)) | ~ sP10), inference(cnf_transformation, [], [f176])).
fof(f1406, plain, (~ spl66_132 | ~ spl66_23), inference(avatar_split_clause, [], [f524, f695, f1403])).
fof(f524, plain, (~ (e2 = op(e2, e2)) | ~ sP10), inference(cnf_transformation, [], [f176])).
fof(f1401, plain, (~ spl66_131 | spl66_53), inference(avatar_split_clause, [], [f517, f823, f1395])).
fof(f517, plain, ((e0 = op(e0, e2)) | ~ sP11), inference(cnf_transformation, [], [f175])).
fof(f175, plain, ((~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | ~ sP11), inference(nnf_transformation, [], [f65])).
fof(f1400, plain, (~ spl66_131 | ~ spl66_63), inference(avatar_split_clause, [], [f518, f865, f1395])).
fof(f518, plain, (~ (op(e0, e0) = e2) | ~ sP11), inference(cnf_transformation, [], [f175])).
fof(f1399, plain, (~ spl66_131 | spl66_3), inference(avatar_split_clause, [], [f519, f610, f1395])).
fof(f519, plain, ((e2 = op(e3, e3)) | ~ sP11), inference(cnf_transformation, [], [f175])).
fof(f1398, plain, (~ spl66_131 | ~ spl66_8), inference(avatar_split_clause, [], [f520, f631, f1395])).
fof(f520, plain, (~ (e3 = op(e3, e2)) | ~ sP11), inference(cnf_transformation, [], [f175])).
fof(f1392, plain, (~ spl66_130 | ~ spl66_64), inference(avatar_split_clause, [], [f514, f869, f1387])).
fof(f514, plain, (~ (op(e0, e0) = e3) | ~ sP12), inference(cnf_transformation, [], [f174])).
fof(f174, plain, ((~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | ~ sP12), inference(nnf_transformation, [], [f66])).
fof(f1391, plain, (~ spl66_130 | spl66_64), inference(avatar_split_clause, [], [f515, f869, f1387])).
fof(f515, plain, ((op(e0, e0) = e3) | ~ sP12), inference(cnf_transformation, [], [f174])).
fof(f1385, plain, (~ spl66_129 | spl66_49), inference(avatar_split_clause, [], [f509, f806, f1379])).
fof(f509, plain, ((e0 = op(e0, e3)) | ~ sP13), inference(cnf_transformation, [], [f173])).
fof(f173, plain, ((~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | ~ sP13), inference(nnf_transformation, [], [f67])).
fof(f1384, plain, (~ spl66_129 | ~ spl66_64), inference(avatar_split_clause, [], [f510, f869, f1379])).
fof(f510, plain, (~ (op(e0, e0) = e3) | ~ sP13), inference(cnf_transformation, [], [f173])).
fof(f1383, plain, (~ spl66_129 | spl66_44), inference(avatar_split_clause, [], [f511, f784, f1379])).
fof(f511, plain, ((e3 = op(e1, e1)) | ~ sP13), inference(cnf_transformation, [], [f173])).
fof(f1382, plain, (~ spl66_129 | ~ spl66_34), inference(avatar_split_clause, [], [f512, f742, f1379])).
fof(f512, plain, (~ (e1 = op(e1, e3)) | ~ sP13), inference(cnf_transformation, [], [f173])).
fof(f1377, plain, (~ spl66_128 | spl66_49), inference(avatar_split_clause, [], [f505, f806, f1371])).
fof(f505, plain, ((e0 = op(e0, e3)) | ~ sP14), inference(cnf_transformation, [], [f172])).
fof(f172, plain, ((~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | ~ sP14), inference(nnf_transformation, [], [f68])).
fof(f1376, plain, (~ spl66_128 | ~ spl66_64), inference(avatar_split_clause, [], [f506, f869, f1371])).
fof(f506, plain, (~ (op(e0, e0) = e3) | ~ sP14), inference(cnf_transformation, [], [f172])).
fof(f1375, plain, (~ spl66_128 | spl66_24), inference(avatar_split_clause, [], [f507, f699, f1371])).
fof(f507, plain, ((e3 = op(e2, e2)) | ~ sP14), inference(cnf_transformation, [], [f172])).
fof(f1374, plain, (~ spl66_128 | ~ spl66_19), inference(avatar_split_clause, [], [f508, f678, f1371])).
fof(f508, plain, (~ (e2 = op(e2, e3)) | ~ sP14), inference(cnf_transformation, [], [f172])).
fof(f1367, plain, (~ spl66_127 | spl66_4), inference(avatar_split_clause, [], [f503, f614, f1363])).
fof(f503, plain, ((e3 = op(e3, e3)) | ~ sP15), inference(cnf_transformation, [], [f171])).
fof(f171, plain, ((~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | ~ sP15), inference(nnf_transformation, [], [f69])).
fof(f1366, plain, (~ spl66_127 | ~ spl66_4), inference(avatar_split_clause, [], [f504, f614, f1363])).
fof(f504, plain, (~ (e3 = op(e3, e3)) | ~ sP15), inference(cnf_transformation, [], [f171])).
fof(f1359, plain, (~ spl66_126 | spl66_61), inference(avatar_split_clause, [], [f499, f857, f1355])).
fof(f499, plain, ((e0 = op(e0, e0)) | ~ sP16), inference(cnf_transformation, [], [f170])).
fof(f170, plain, ((~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | ~ sP16), inference(nnf_transformation, [], [f70])).
fof(f1358, plain, (~ spl66_126 | ~ spl66_61), inference(avatar_split_clause, [], [f500, f857, f1355])).
fof(f500, plain, (~ (e0 = op(e0, e0)) | ~ sP16), inference(cnf_transformation, [], [f170])).
fof(f1353, plain, (~ spl66_125 | spl66_46), inference(avatar_split_clause, [], [f493, f793, f1347])).
fof(f493, plain, ((e1 = op(e1, e0)) | ~ sP17), inference(cnf_transformation, [], [f169])).
fof(f169, plain, ((~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | ~ sP17), inference(nnf_transformation, [], [f71])).
fof(f1351, plain, (~ spl66_125 | spl66_41), inference(avatar_split_clause, [], [f495, f772, f1347])).
fof(f495, plain, ((e0 = op(e1, e1)) | ~ sP17), inference(cnf_transformation, [], [f169])).
fof(f1350, plain, (~ spl66_125 | ~ spl66_46), inference(avatar_split_clause, [], [f496, f793, f1347])).
fof(f496, plain, (~ (e1 = op(e1, e0)) | ~ sP17), inference(cnf_transformation, [], [f169])).
fof(f1345, plain, (~ spl66_124 | spl66_46), inference(avatar_split_clause, [], [f489, f793, f1339])).
fof(f489, plain, ((e1 = op(e1, e0)) | ~ sP18), inference(cnf_transformation, [], [f168])).
fof(f168, plain, ((~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | ~ sP18), inference(nnf_transformation, [], [f72])).
fof(f1344, plain, (~ spl66_124 | ~ spl66_41), inference(avatar_split_clause, [], [f490, f772, f1339])).
fof(f490, plain, (~ (e0 = op(e1, e1)) | ~ sP18), inference(cnf_transformation, [], [f168])).
fof(f1343, plain, (~ spl66_124 | spl66_21), inference(avatar_split_clause, [], [f491, f687, f1339])).
fof(f491, plain, ((e0 = op(e2, e2)) | ~ sP18), inference(cnf_transformation, [], [f168])).
fof(f1342, plain, (~ spl66_124 | ~ spl66_31), inference(avatar_split_clause, [], [f492, f729, f1339])).
fof(f492, plain, (~ (e2 = op(e2, e0)) | ~ sP18), inference(cnf_transformation, [], [f168])).
fof(f1337, plain, (~ spl66_123 | spl66_46), inference(avatar_split_clause, [], [f485, f793, f1331])).
fof(f485, plain, ((e1 = op(e1, e0)) | ~ sP19), inference(cnf_transformation, [], [f167])).
fof(f167, plain, ((~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | ~ sP19), inference(nnf_transformation, [], [f73])).
fof(f1336, plain, (~ spl66_123 | ~ spl66_41), inference(avatar_split_clause, [], [f486, f772, f1331])).
fof(f486, plain, (~ (e0 = op(e1, e1)) | ~ sP19), inference(cnf_transformation, [], [f167])).
fof(f1335, plain, (~ spl66_123 | spl66_1), inference(avatar_split_clause, [], [f487, f602, f1331])).
fof(f487, plain, ((e0 = op(e3, e3)) | ~ sP19), inference(cnf_transformation, [], [f167])).
fof(f1334, plain, (~ spl66_123 | ~ spl66_16), inference(avatar_split_clause, [], [f488, f665, f1331])).
fof(f488, plain, (~ (e3 = op(e3, e0)) | ~ sP19), inference(cnf_transformation, [], [f167])).
fof(f1329, plain, (~ spl66_122 | spl66_42), inference(avatar_split_clause, [], [f481, f776, f1323])).
fof(f481, plain, ((e1 = op(e1, e1)) | ~ sP20), inference(cnf_transformation, [], [f166])).
fof(f166, plain, ((~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP20), inference(nnf_transformation, [], [f74])).
fof(f1328, plain, (~ spl66_122 | ~ spl66_42), inference(avatar_split_clause, [], [f482, f776, f1323])).
fof(f482, plain, (~ (e1 = op(e1, e1)) | ~ sP20), inference(cnf_transformation, [], [f166])).
fof(f1321, plain, (~ spl66_121 | spl66_42), inference(avatar_split_clause, [], [f477, f776, f1315])).
fof(f477, plain, ((e1 = op(e1, e1)) | ~ sP21), inference(cnf_transformation, [], [f165])).
fof(f165, plain, ((~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP21), inference(nnf_transformation, [], [f75])).
fof(f1320, plain, (~ spl66_121 | ~ spl66_42), inference(avatar_split_clause, [], [f478, f776, f1315])).
fof(f478, plain, (~ (e1 = op(e1, e1)) | ~ sP21), inference(cnf_transformation, [], [f165])).
fof(f1313, plain, (~ spl66_120 | spl66_42), inference(avatar_split_clause, [], [f473, f776, f1307])).
fof(f473, plain, ((e1 = op(e1, e1)) | ~ sP22), inference(cnf_transformation, [], [f164])).
fof(f164, plain, ((~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP22), inference(nnf_transformation, [], [f76])).
fof(f1312, plain, (~ spl66_120 | ~ spl66_42), inference(avatar_split_clause, [], [f474, f776, f1307])).
fof(f474, plain, (~ (e1 = op(e1, e1)) | ~ sP22), inference(cnf_transformation, [], [f164])).
fof(f1305, plain, (~ spl66_119 | spl66_42), inference(avatar_split_clause, [], [f469, f776, f1299])).
fof(f469, plain, ((e1 = op(e1, e1)) | ~ sP23), inference(cnf_transformation, [], [f163])).
fof(f163, plain, ((~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP23), inference(nnf_transformation, [], [f77])).
fof(f1304, plain, (~ spl66_119 | ~ spl66_42), inference(avatar_split_clause, [], [f470, f776, f1299])).
fof(f470, plain, (~ (e1 = op(e1, e1)) | ~ sP23), inference(cnf_transformation, [], [f163])).
fof(f1297, plain, (~ spl66_118 | spl66_38), inference(avatar_split_clause, [], [f465, f759, f1291])).
fof(f465, plain, ((e1 = op(e1, e2)) | ~ sP24), inference(cnf_transformation, [], [f162])).
fof(f162, plain, ((~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | ~ sP24), inference(nnf_transformation, [], [f78])).
fof(f1296, plain, (~ spl66_118 | ~ spl66_43), inference(avatar_split_clause, [], [f466, f780, f1291])).
fof(f466, plain, (~ (e2 = op(e1, e1)) | ~ sP24), inference(cnf_transformation, [], [f162])).
fof(f1295, plain, (~ spl66_118 | spl66_63), inference(avatar_split_clause, [], [f467, f865, f1291])).
fof(f467, plain, ((op(e0, e0) = e2) | ~ sP24), inference(cnf_transformation, [], [f162])).
fof(f1294, plain, (~ spl66_118 | ~ spl66_53), inference(avatar_split_clause, [], [f468, f823, f1291])).
fof(f468, plain, (~ (e0 = op(e0, e2)) | ~ sP24), inference(cnf_transformation, [], [f162])).
fof(f1288, plain, (~ spl66_117 | ~ spl66_43), inference(avatar_split_clause, [], [f462, f780, f1283])).
fof(f462, plain, (~ (e2 = op(e1, e1)) | ~ sP25), inference(cnf_transformation, [], [f161])).
fof(f161, plain, ((~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | ~ sP25), inference(nnf_transformation, [], [f79])).
fof(f1287, plain, (~ spl66_117 | spl66_43), inference(avatar_split_clause, [], [f463, f780, f1283])).
fof(f463, plain, ((e2 = op(e1, e1)) | ~ sP25), inference(cnf_transformation, [], [f161])).
fof(f1281, plain, (~ spl66_116 | spl66_38), inference(avatar_split_clause, [], [f457, f759, f1275])).
fof(f457, plain, ((e1 = op(e1, e2)) | ~ sP26), inference(cnf_transformation, [], [f160])).
fof(f160, plain, ((~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | ~ sP26), inference(nnf_transformation, [], [f80])).
fof(f1280, plain, (~ spl66_116 | ~ spl66_43), inference(avatar_split_clause, [], [f458, f780, f1275])).
fof(f458, plain, (~ (e2 = op(e1, e1)) | ~ sP26), inference(cnf_transformation, [], [f160])).
fof(f1279, plain, (~ spl66_116 | spl66_23), inference(avatar_split_clause, [], [f459, f695, f1275])).
fof(f459, plain, ((e2 = op(e2, e2)) | ~ sP26), inference(cnf_transformation, [], [f160])).
fof(f1278, plain, (~ spl66_116 | ~ spl66_23), inference(avatar_split_clause, [], [f460, f695, f1275])).
fof(f460, plain, (~ (e2 = op(e2, e2)) | ~ sP26), inference(cnf_transformation, [], [f160])).
fof(f1273, plain, (~ spl66_115 | spl66_38), inference(avatar_split_clause, [], [f453, f759, f1267])).
fof(f453, plain, ((e1 = op(e1, e2)) | ~ sP27), inference(cnf_transformation, [], [f159])).
fof(f159, plain, ((~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | ~ sP27), inference(nnf_transformation, [], [f81])).
fof(f1272, plain, (~ spl66_115 | ~ spl66_43), inference(avatar_split_clause, [], [f454, f780, f1267])).
fof(f454, plain, (~ (e2 = op(e1, e1)) | ~ sP27), inference(cnf_transformation, [], [f159])).
fof(f1271, plain, (~ spl66_115 | spl66_3), inference(avatar_split_clause, [], [f455, f610, f1267])).
fof(f455, plain, ((e2 = op(e3, e3)) | ~ sP27), inference(cnf_transformation, [], [f159])).
fof(f1270, plain, (~ spl66_115 | ~ spl66_8), inference(avatar_split_clause, [], [f456, f631, f1267])).
fof(f456, plain, (~ (e3 = op(e3, e2)) | ~ sP27), inference(cnf_transformation, [], [f159])).
fof(f1265, plain, (~ spl66_114 | spl66_34), inference(avatar_split_clause, [], [f449, f742, f1259])).
fof(f449, plain, ((e1 = op(e1, e3)) | ~ sP28), inference(cnf_transformation, [], [f158])).
fof(f158, plain, ((~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | ~ sP28), inference(nnf_transformation, [], [f82])).
fof(f1264, plain, (~ spl66_114 | ~ spl66_44), inference(avatar_split_clause, [], [f450, f784, f1259])).
fof(f450, plain, (~ (e3 = op(e1, e1)) | ~ sP28), inference(cnf_transformation, [], [f158])).
fof(f1263, plain, (~ spl66_114 | spl66_64), inference(avatar_split_clause, [], [f451, f869, f1259])).
fof(f451, plain, ((op(e0, e0) = e3) | ~ sP28), inference(cnf_transformation, [], [f158])).
fof(f1262, plain, (~ spl66_114 | ~ spl66_49), inference(avatar_split_clause, [], [f452, f806, f1259])).
fof(f452, plain, (~ (e0 = op(e0, e3)) | ~ sP28), inference(cnf_transformation, [], [f158])).
fof(f1256, plain, (~ spl66_113 | ~ spl66_44), inference(avatar_split_clause, [], [f446, f784, f1251])).
fof(f446, plain, (~ (e3 = op(e1, e1)) | ~ sP29), inference(cnf_transformation, [], [f157])).
fof(f157, plain, ((~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | ~ sP29), inference(nnf_transformation, [], [f83])).
fof(f1255, plain, (~ spl66_113 | spl66_44), inference(avatar_split_clause, [], [f447, f784, f1251])).
fof(f447, plain, ((e3 = op(e1, e1)) | ~ sP29), inference(cnf_transformation, [], [f157])).
fof(f1249, plain, (~ spl66_112 | spl66_34), inference(avatar_split_clause, [], [f441, f742, f1243])).
fof(f441, plain, ((e1 = op(e1, e3)) | ~ sP30), inference(cnf_transformation, [], [f156])).
fof(f156, plain, ((~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | ~ sP30), inference(nnf_transformation, [], [f84])).
fof(f1248, plain, (~ spl66_112 | ~ spl66_44), inference(avatar_split_clause, [], [f442, f784, f1243])).
fof(f442, plain, (~ (e3 = op(e1, e1)) | ~ sP30), inference(cnf_transformation, [], [f156])).
fof(f1247, plain, (~ spl66_112 | spl66_24), inference(avatar_split_clause, [], [f443, f699, f1243])).
fof(f443, plain, ((e3 = op(e2, e2)) | ~ sP30), inference(cnf_transformation, [], [f156])).
fof(f1246, plain, (~ spl66_112 | ~ spl66_19), inference(avatar_split_clause, [], [f444, f678, f1243])).
fof(f444, plain, (~ (e2 = op(e2, e3)) | ~ sP30), inference(cnf_transformation, [], [f156])).
fof(f1241, plain, (~ spl66_111 | spl66_34), inference(avatar_split_clause, [], [f437, f742, f1235])).
fof(f437, plain, ((e1 = op(e1, e3)) | ~ sP31), inference(cnf_transformation, [], [f155])).
fof(f155, plain, ((~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | ~ sP31), inference(nnf_transformation, [], [f85])).
fof(f1239, plain, (~ spl66_111 | spl66_4), inference(avatar_split_clause, [], [f439, f614, f1235])).
fof(f439, plain, ((e3 = op(e3, e3)) | ~ sP31), inference(cnf_transformation, [], [f155])).
fof(f1238, plain, (~ spl66_111 | ~ spl66_4), inference(avatar_split_clause, [], [f440, f614, f1235])).
fof(f440, plain, (~ (e3 = op(e3, e3)) | ~ sP31), inference(cnf_transformation, [], [f155])).
fof(f1231, plain, (~ spl66_110 | spl66_61), inference(avatar_split_clause, [], [f435, f857, f1227])).
fof(f435, plain, ((e0 = op(e0, e0)) | ~ sP32), inference(cnf_transformation, [], [f154])).
fof(f154, plain, ((~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | ~ sP32), inference(nnf_transformation, [], [f86])).
fof(f1230, plain, (~ spl66_110 | ~ spl66_61), inference(avatar_split_clause, [], [f436, f857, f1227])).
fof(f436, plain, (~ (e0 = op(e0, e0)) | ~ sP32), inference(cnf_transformation, [], [f154])).
fof(f1225, plain, (~ spl66_109 | spl66_31), inference(avatar_split_clause, [], [f429, f729, f1219])).
fof(f429, plain, ((e2 = op(e2, e0)) | ~ sP33), inference(cnf_transformation, [], [f153])).
fof(f153, plain, ((~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | ~ sP33), inference(nnf_transformation, [], [f87])).
fof(f1224, plain, (~ spl66_109 | ~ spl66_21), inference(avatar_split_clause, [], [f430, f687, f1219])).
fof(f430, plain, (~ (e0 = op(e2, e2)) | ~ sP33), inference(cnf_transformation, [], [f153])).
fof(f1223, plain, (~ spl66_109 | spl66_41), inference(avatar_split_clause, [], [f431, f772, f1219])).
fof(f431, plain, ((e0 = op(e1, e1)) | ~ sP33), inference(cnf_transformation, [], [f153])).
fof(f1222, plain, (~ spl66_109 | ~ spl66_46), inference(avatar_split_clause, [], [f432, f793, f1219])).
fof(f432, plain, (~ (e1 = op(e1, e0)) | ~ sP33), inference(cnf_transformation, [], [f153])).
fof(f1216, plain, (~ spl66_108 | ~ spl66_21), inference(avatar_split_clause, [], [f426, f687, f1211])).
fof(f426, plain, (~ (e0 = op(e2, e2)) | ~ sP34), inference(cnf_transformation, [], [f152])).
fof(f152, plain, ((~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | ~ sP34), inference(nnf_transformation, [], [f88])).
fof(f1215, plain, (~ spl66_108 | spl66_21), inference(avatar_split_clause, [], [f427, f687, f1211])).
fof(f427, plain, ((e0 = op(e2, e2)) | ~ sP34), inference(cnf_transformation, [], [f152])).
fof(f1209, plain, (~ spl66_107 | spl66_31), inference(avatar_split_clause, [], [f421, f729, f1203])).
fof(f421, plain, ((e2 = op(e2, e0)) | ~ sP35), inference(cnf_transformation, [], [f151])).
fof(f151, plain, ((~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | ~ sP35), inference(nnf_transformation, [], [f89])).
fof(f1208, plain, (~ spl66_107 | ~ spl66_21), inference(avatar_split_clause, [], [f422, f687, f1203])).
fof(f422, plain, (~ (e0 = op(e2, e2)) | ~ sP35), inference(cnf_transformation, [], [f151])).
fof(f1207, plain, (~ spl66_107 | spl66_1), inference(avatar_split_clause, [], [f423, f602, f1203])).
fof(f423, plain, ((e0 = op(e3, e3)) | ~ sP35), inference(cnf_transformation, [], [f151])).
fof(f1206, plain, (~ spl66_107 | ~ spl66_16), inference(avatar_split_clause, [], [f424, f665, f1203])).
fof(f424, plain, (~ (e3 = op(e3, e0)) | ~ sP35), inference(cnf_transformation, [], [f151])).
fof(f1201, plain, (~ spl66_106 | spl66_27), inference(avatar_split_clause, [], [f417, f712, f1195])).
fof(f417, plain, ((e2 = op(e2, e1)) | ~ sP36), inference(cnf_transformation, [], [f150])).
fof(f150, plain, ((~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | ~ sP36), inference(nnf_transformation, [], [f90])).
fof(f1200, plain, (~ spl66_106 | ~ spl66_22), inference(avatar_split_clause, [], [f418, f691, f1195])).
fof(f418, plain, (~ (e1 = op(e2, e2)) | ~ sP36), inference(cnf_transformation, [], [f150])).
fof(f1199, plain, (~ spl66_106 | spl66_62), inference(avatar_split_clause, [], [f419, f861, f1195])).
fof(f419, plain, ((op(e0, e0) = e1) | ~ sP36), inference(cnf_transformation, [], [f150])).
fof(f1198, plain, (~ spl66_106 | ~ spl66_57), inference(avatar_split_clause, [], [f420, f840, f1195])).
fof(f420, plain, (~ (e0 = op(e0, e1)) | ~ sP36), inference(cnf_transformation, [], [f150])).
fof(f1191, plain, (~ spl66_105 | spl66_42), inference(avatar_split_clause, [], [f415, f776, f1187])).
fof(f415, plain, ((e1 = op(e1, e1)) | ~ sP37), inference(cnf_transformation, [], [f149])).
fof(f149, plain, ((~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | ~ sP37), inference(nnf_transformation, [], [f91])).
fof(f1190, plain, (~ spl66_105 | ~ spl66_42), inference(avatar_split_clause, [], [f416, f776, f1187])).
fof(f416, plain, (~ (e1 = op(e1, e1)) | ~ sP37), inference(cnf_transformation, [], [f149])).
fof(f1184, plain, (~ spl66_104 | ~ spl66_22), inference(avatar_split_clause, [], [f410, f691, f1179])).
fof(f410, plain, (~ (e1 = op(e2, e2)) | ~ sP38), inference(cnf_transformation, [], [f148])).
fof(f148, plain, ((~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | ~ sP38), inference(nnf_transformation, [], [f92])).
fof(f1183, plain, (~ spl66_104 | spl66_22), inference(avatar_split_clause, [], [f411, f691, f1179])).
fof(f411, plain, ((e1 = op(e2, e2)) | ~ sP38), inference(cnf_transformation, [], [f148])).
fof(f1177, plain, (~ spl66_103 | spl66_27), inference(avatar_split_clause, [], [f405, f712, f1171])).
fof(f405, plain, ((e2 = op(e2, e1)) | ~ sP39), inference(cnf_transformation, [], [f147])).
fof(f147, plain, ((~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | ~ sP39), inference(nnf_transformation, [], [f93])).
fof(f1176, plain, (~ spl66_103 | ~ spl66_22), inference(avatar_split_clause, [], [f406, f691, f1171])).
fof(f406, plain, (~ (e1 = op(e2, e2)) | ~ sP39), inference(cnf_transformation, [], [f147])).
fof(f1175, plain, (~ spl66_103 | spl66_2), inference(avatar_split_clause, [], [f407, f606, f1171])).
fof(f407, plain, ((e1 = op(e3, e3)) | ~ sP39), inference(cnf_transformation, [], [f147])).
fof(f1174, plain, (~ spl66_103 | ~ spl66_12), inference(avatar_split_clause, [], [f408, f648, f1171])).
fof(f408, plain, (~ (e3 = op(e3, e1)) | ~ sP39), inference(cnf_transformation, [], [f147])).
fof(f1169, plain, (~ spl66_102 | spl66_23), inference(avatar_split_clause, [], [f401, f695, f1163])).
fof(f401, plain, ((e2 = op(e2, e2)) | ~ sP40), inference(cnf_transformation, [], [f146])).
fof(f146, plain, ((~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP40), inference(nnf_transformation, [], [f94])).
fof(f1168, plain, (~ spl66_102 | ~ spl66_23), inference(avatar_split_clause, [], [f402, f695, f1163])).
fof(f402, plain, (~ (e2 = op(e2, e2)) | ~ sP40), inference(cnf_transformation, [], [f146])).
fof(f1161, plain, (~ spl66_101 | spl66_23), inference(avatar_split_clause, [], [f397, f695, f1155])).
fof(f397, plain, ((e2 = op(e2, e2)) | ~ sP41), inference(cnf_transformation, [], [f145])).
fof(f145, plain, ((~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP41), inference(nnf_transformation, [], [f95])).
fof(f1160, plain, (~ spl66_101 | ~ spl66_23), inference(avatar_split_clause, [], [f398, f695, f1155])).
fof(f398, plain, (~ (e2 = op(e2, e2)) | ~ sP41), inference(cnf_transformation, [], [f145])).
fof(f1153, plain, (~ spl66_100 | spl66_23), inference(avatar_split_clause, [], [f393, f695, f1147])).
fof(f393, plain, ((e2 = op(e2, e2)) | ~ sP42), inference(cnf_transformation, [], [f144])).
fof(f144, plain, ((~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP42), inference(nnf_transformation, [], [f96])).
fof(f1152, plain, (~ spl66_100 | ~ spl66_23), inference(avatar_split_clause, [], [f394, f695, f1147])).
fof(f394, plain, (~ (e2 = op(e2, e2)) | ~ sP42), inference(cnf_transformation, [], [f144])).
fof(f1145, plain, (~ spl66_99 | spl66_23), inference(avatar_split_clause, [], [f389, f695, f1139])).
fof(f389, plain, ((e2 = op(e2, e2)) | ~ sP43), inference(cnf_transformation, [], [f143])).
fof(f143, plain, ((~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP43), inference(nnf_transformation, [], [f97])).
fof(f1144, plain, (~ spl66_99 | ~ spl66_23), inference(avatar_split_clause, [], [f390, f695, f1139])).
fof(f390, plain, (~ (e2 = op(e2, e2)) | ~ sP43), inference(cnf_transformation, [], [f143])).
fof(f1143, plain, (~ spl66_99 | spl66_3), inference(avatar_split_clause, [], [f391, f610, f1139])).
fof(f391, plain, ((e2 = op(e3, e3)) | ~ sP43), inference(cnf_transformation, [], [f143])).
fof(f1137, plain, (~ spl66_98 | spl66_19), inference(avatar_split_clause, [], [f385, f678, f1131])).
fof(f385, plain, ((e2 = op(e2, e3)) | ~ sP44), inference(cnf_transformation, [], [f142])).
fof(f142, plain, ((~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | ~ sP44), inference(nnf_transformation, [], [f98])).
fof(f1136, plain, (~ spl66_98 | ~ spl66_24), inference(avatar_split_clause, [], [f386, f699, f1131])).
fof(f386, plain, (~ (e3 = op(e2, e2)) | ~ sP44), inference(cnf_transformation, [], [f142])).
fof(f1135, plain, (~ spl66_98 | spl66_64), inference(avatar_split_clause, [], [f387, f869, f1131])).
fof(f387, plain, ((op(e0, e0) = e3) | ~ sP44), inference(cnf_transformation, [], [f142])).
fof(f1134, plain, (~ spl66_98 | ~ spl66_49), inference(avatar_split_clause, [], [f388, f806, f1131])).
fof(f388, plain, (~ (e0 = op(e0, e3)) | ~ sP44), inference(cnf_transformation, [], [f142])).
fof(f1129, plain, (~ spl66_97 | spl66_19), inference(avatar_split_clause, [], [f381, f678, f1123])).
fof(f381, plain, ((e2 = op(e2, e3)) | ~ sP45), inference(cnf_transformation, [], [f141])).
fof(f141, plain, ((~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | ~ sP45), inference(nnf_transformation, [], [f99])).
fof(f1128, plain, (~ spl66_97 | ~ spl66_24), inference(avatar_split_clause, [], [f382, f699, f1123])).
fof(f382, plain, (~ (e3 = op(e2, e2)) | ~ sP45), inference(cnf_transformation, [], [f141])).
fof(f1127, plain, (~ spl66_97 | spl66_44), inference(avatar_split_clause, [], [f383, f784, f1123])).
fof(f383, plain, ((e3 = op(e1, e1)) | ~ sP45), inference(cnf_transformation, [], [f141])).
fof(f1126, plain, (~ spl66_97 | ~ spl66_34), inference(avatar_split_clause, [], [f384, f742, f1123])).
fof(f384, plain, (~ (e1 = op(e1, e3)) | ~ sP45), inference(cnf_transformation, [], [f141])).
fof(f1120, plain, (~ spl66_96 | ~ spl66_24), inference(avatar_split_clause, [], [f378, f699, f1115])).
fof(f378, plain, (~ (e3 = op(e2, e2)) | ~ sP46), inference(cnf_transformation, [], [f140])).
fof(f140, plain, ((~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | ~ sP46), inference(nnf_transformation, [], [f100])).
fof(f1119, plain, (~ spl66_96 | spl66_24), inference(avatar_split_clause, [], [f379, f699, f1115])).
fof(f379, plain, ((e3 = op(e2, e2)) | ~ sP46), inference(cnf_transformation, [], [f140])).
fof(f1111, plain, (~ spl66_95 | spl66_4), inference(avatar_split_clause, [], [f375, f614, f1107])).
fof(f375, plain, ((e3 = op(e3, e3)) | ~ sP47), inference(cnf_transformation, [], [f139])).
fof(f139, plain, ((~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | ~ sP47), inference(nnf_transformation, [], [f101])).
fof(f1110, plain, (~ spl66_95 | ~ spl66_4), inference(avatar_split_clause, [], [f376, f614, f1107])).
fof(f376, plain, (~ (e3 = op(e3, e3)) | ~ sP47), inference(cnf_transformation, [], [f139])).
fof(f1105, plain, (~ spl66_94 | spl66_16), inference(avatar_split_clause, [], [f369, f665, f1099])).
fof(f369, plain, ((e3 = op(e3, e0)) | ~ sP48), inference(cnf_transformation, [], [f138])).
fof(f138, plain, ((~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | ~ sP48), inference(nnf_transformation, [], [f102])).
fof(f1104, plain, (~ spl66_94 | ~ spl66_1), inference(avatar_split_clause, [], [f370, f602, f1099])).
fof(f370, plain, (~ (e0 = op(e3, e3)) | ~ sP48), inference(cnf_transformation, [], [f138])).
fof(f1103, plain, (~ spl66_94 | spl66_61), inference(avatar_split_clause, [], [f371, f857, f1099])).
fof(f371, plain, ((e0 = op(e0, e0)) | ~ sP48), inference(cnf_transformation, [], [f138])).
fof(f1102, plain, (~ spl66_94 | ~ spl66_61), inference(avatar_split_clause, [], [f372, f857, f1099])).
fof(f372, plain, (~ (e0 = op(e0, e0)) | ~ sP48), inference(cnf_transformation, [], [f138])).
fof(f1097, plain, (~ spl66_93 | spl66_16), inference(avatar_split_clause, [], [f365, f665, f1091])).
fof(f365, plain, ((e3 = op(e3, e0)) | ~ sP49), inference(cnf_transformation, [], [f137])).
fof(f137, plain, ((~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | ~ sP49), inference(nnf_transformation, [], [f103])).
fof(f1096, plain, (~ spl66_93 | ~ spl66_1), inference(avatar_split_clause, [], [f366, f602, f1091])).
fof(f366, plain, (~ (e0 = op(e3, e3)) | ~ sP49), inference(cnf_transformation, [], [f137])).
fof(f1095, plain, (~ spl66_93 | spl66_41), inference(avatar_split_clause, [], [f367, f772, f1091])).
fof(f367, plain, ((e0 = op(e1, e1)) | ~ sP49), inference(cnf_transformation, [], [f137])).
fof(f1094, plain, (~ spl66_93 | ~ spl66_46), inference(avatar_split_clause, [], [f368, f793, f1091])).
fof(f368, plain, (~ (e1 = op(e1, e0)) | ~ sP49), inference(cnf_transformation, [], [f137])).
fof(f1089, plain, (~ spl66_92 | spl66_16), inference(avatar_split_clause, [], [f361, f665, f1083])).
fof(f361, plain, ((e3 = op(e3, e0)) | ~ sP50), inference(cnf_transformation, [], [f136])).
fof(f136, plain, ((~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | ~ sP50), inference(nnf_transformation, [], [f104])).
fof(f1088, plain, (~ spl66_92 | ~ spl66_1), inference(avatar_split_clause, [], [f362, f602, f1083])).
fof(f362, plain, (~ (e0 = op(e3, e3)) | ~ sP50), inference(cnf_transformation, [], [f136])).
fof(f1087, plain, (~ spl66_92 | spl66_21), inference(avatar_split_clause, [], [f363, f687, f1083])).
fof(f363, plain, ((e0 = op(e2, e2)) | ~ sP50), inference(cnf_transformation, [], [f136])).
fof(f1086, plain, (~ spl66_92 | ~ spl66_31), inference(avatar_split_clause, [], [f364, f729, f1083])).
fof(f364, plain, (~ (e2 = op(e2, e0)) | ~ sP50), inference(cnf_transformation, [], [f136])).
fof(f1081, plain, (~ spl66_91 | spl66_16), inference(avatar_split_clause, [], [f357, f665, f1075])).
fof(f357, plain, ((e3 = op(e3, e0)) | ~ sP51), inference(cnf_transformation, [], [f135])).
fof(f135, plain, ((~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | ~ sP51), inference(nnf_transformation, [], [f105])).
fof(f1080, plain, (~ spl66_91 | ~ spl66_1), inference(avatar_split_clause, [], [f358, f602, f1075])).
fof(f358, plain, (~ (e0 = op(e3, e3)) | ~ sP51), inference(cnf_transformation, [], [f135])).
fof(f1079, plain, (~ spl66_91 | spl66_1), inference(avatar_split_clause, [], [f359, f602, f1075])).
fof(f359, plain, ((e0 = op(e3, e3)) | ~ sP51), inference(cnf_transformation, [], [f135])).
fof(f1078, plain, (~ spl66_91 | ~ spl66_16), inference(avatar_split_clause, [], [f360, f665, f1075])).
fof(f360, plain, (~ (e3 = op(e3, e0)) | ~ sP51), inference(cnf_transformation, [], [f135])).
fof(f1073, plain, (~ spl66_90 | spl66_12), inference(avatar_split_clause, [], [f353, f648, f1067])).
fof(f353, plain, ((e3 = op(e3, e1)) | ~ sP52), inference(cnf_transformation, [], [f134])).
fof(f134, plain, ((~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | ~ sP52), inference(nnf_transformation, [], [f106])).
fof(f1072, plain, (~ spl66_90 | ~ spl66_2), inference(avatar_split_clause, [], [f354, f606, f1067])).
fof(f354, plain, (~ (e1 = op(e3, e3)) | ~ sP52), inference(cnf_transformation, [], [f134])).
fof(f1071, plain, (~ spl66_90 | spl66_62), inference(avatar_split_clause, [], [f355, f861, f1067])).
fof(f355, plain, ((op(e0, e0) = e1) | ~ sP52), inference(cnf_transformation, [], [f134])).
fof(f1070, plain, (~ spl66_90 | ~ spl66_57), inference(avatar_split_clause, [], [f356, f840, f1067])).
fof(f356, plain, (~ (e0 = op(e0, e1)) | ~ sP52), inference(cnf_transformation, [], [f134])).
fof(f1063, plain, (~ spl66_89 | spl66_42), inference(avatar_split_clause, [], [f351, f776, f1059])).
fof(f351, plain, ((e1 = op(e1, e1)) | ~ sP53), inference(cnf_transformation, [], [f133])).
fof(f133, plain, ((~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | ~ sP53), inference(nnf_transformation, [], [f107])).
fof(f1062, plain, (~ spl66_89 | ~ spl66_42), inference(avatar_split_clause, [], [f352, f776, f1059])).
fof(f352, plain, (~ (e1 = op(e1, e1)) | ~ sP53), inference(cnf_transformation, [], [f133])).
fof(f1057, plain, (~ spl66_88 | spl66_12), inference(avatar_split_clause, [], [f345, f648, f1051])).
fof(f345, plain, ((e3 = op(e3, e1)) | ~ sP54), inference(cnf_transformation, [], [f132])).
fof(f132, plain, ((~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | ~ sP54), inference(nnf_transformation, [], [f108])).
fof(f1056, plain, (~ spl66_88 | ~ spl66_2), inference(avatar_split_clause, [], [f346, f606, f1051])).
fof(f346, plain, (~ (e1 = op(e3, e3)) | ~ sP54), inference(cnf_transformation, [], [f132])).
fof(f1055, plain, (~ spl66_88 | spl66_22), inference(avatar_split_clause, [], [f347, f691, f1051])).
fof(f347, plain, ((e1 = op(e2, e2)) | ~ sP54), inference(cnf_transformation, [], [f132])).
fof(f1054, plain, (~ spl66_88 | ~ spl66_27), inference(avatar_split_clause, [], [f348, f712, f1051])).
fof(f348, plain, (~ (e2 = op(e2, e1)) | ~ sP54), inference(cnf_transformation, [], [f132])).
fof(f1049, plain, (~ spl66_87 | spl66_12), inference(avatar_split_clause, [], [f341, f648, f1043])).
fof(f341, plain, ((e3 = op(e3, e1)) | ~ sP55), inference(cnf_transformation, [], [f131])).
fof(f131, plain, ((~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | ~ sP55), inference(nnf_transformation, [], [f109])).
fof(f1048, plain, (~ spl66_87 | ~ spl66_2), inference(avatar_split_clause, [], [f342, f606, f1043])).
fof(f342, plain, (~ (e1 = op(e3, e3)) | ~ sP55), inference(cnf_transformation, [], [f131])).
fof(f1047, plain, (~ spl66_87 | spl66_2), inference(avatar_split_clause, [], [f343, f606, f1043])).
fof(f343, plain, ((e1 = op(e3, e3)) | ~ sP55), inference(cnf_transformation, [], [f131])).
fof(f1041, plain, (~ spl66_86 | spl66_8), inference(avatar_split_clause, [], [f337, f631, f1035])).
fof(f337, plain, ((e3 = op(e3, e2)) | ~ sP56), inference(cnf_transformation, [], [f130])).
fof(f130, plain, ((~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | ~ sP56), inference(nnf_transformation, [], [f110])).
fof(f1040, plain, (~ spl66_86 | ~ spl66_3), inference(avatar_split_clause, [], [f338, f610, f1035])).
fof(f338, plain, (~ (e2 = op(e3, e3)) | ~ sP56), inference(cnf_transformation, [], [f130])).
fof(f1039, plain, (~ spl66_86 | spl66_63), inference(avatar_split_clause, [], [f339, f865, f1035])).
fof(f339, plain, ((op(e0, e0) = e2) | ~ sP56), inference(cnf_transformation, [], [f130])).
fof(f1038, plain, (~ spl66_86 | ~ spl66_53), inference(avatar_split_clause, [], [f340, f823, f1035])).
fof(f340, plain, (~ (e0 = op(e0, e2)) | ~ sP56), inference(cnf_transformation, [], [f130])).
fof(f1033, plain, (~ spl66_85 | spl66_8), inference(avatar_split_clause, [], [f333, f631, f1027])).
fof(f333, plain, ((e3 = op(e3, e2)) | ~ sP57), inference(cnf_transformation, [], [f129])).
fof(f129, plain, ((~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | ~ sP57), inference(nnf_transformation, [], [f111])).
fof(f1032, plain, (~ spl66_85 | ~ spl66_3), inference(avatar_split_clause, [], [f334, f610, f1027])).
fof(f334, plain, (~ (e2 = op(e3, e3)) | ~ sP57), inference(cnf_transformation, [], [f129])).
fof(f1031, plain, (~ spl66_85 | spl66_43), inference(avatar_split_clause, [], [f335, f780, f1027])).
fof(f335, plain, ((e2 = op(e1, e1)) | ~ sP57), inference(cnf_transformation, [], [f129])).
fof(f1030, plain, (~ spl66_85 | ~ spl66_38), inference(avatar_split_clause, [], [f336, f759, f1027])).
fof(f336, plain, (~ (e1 = op(e1, e2)) | ~ sP57), inference(cnf_transformation, [], [f129])).
fof(f1023, plain, (~ spl66_84 | spl66_23), inference(avatar_split_clause, [], [f331, f695, f1019])).
fof(f331, plain, ((e2 = op(e2, e2)) | ~ sP58), inference(cnf_transformation, [], [f128])).
fof(f128, plain, ((~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | ~ sP58), inference(nnf_transformation, [], [f112])).
fof(f1022, plain, (~ spl66_84 | ~ spl66_23), inference(avatar_split_clause, [], [f332, f695, f1019])).
fof(f332, plain, (~ (e2 = op(e2, e2)) | ~ sP58), inference(cnf_transformation, [], [f128])).
fof(f1016, plain, (~ spl66_83 | ~ spl66_3), inference(avatar_split_clause, [], [f326, f610, f1011])).
fof(f326, plain, (~ (e2 = op(e3, e3)) | ~ sP59), inference(cnf_transformation, [], [f127])).
fof(f127, plain, ((~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | ~ sP59), inference(nnf_transformation, [], [f113])).
fof(f1015, plain, (~ spl66_83 | spl66_3), inference(avatar_split_clause, [], [f327, f610, f1011])).
fof(f327, plain, ((e2 = op(e3, e3)) | ~ sP59), inference(cnf_transformation, [], [f127])).
fof(f1009, plain, (~ spl66_82 | spl66_4), inference(avatar_split_clause, [], [f321, f614, f1003])).
fof(f321, plain, ((e3 = op(e3, e3)) | ~ sP60), inference(cnf_transformation, [], [f126])).
fof(f126, plain, ((~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ~ sP60), inference(nnf_transformation, [], [f114])).
fof(f1008, plain, (~ spl66_82 | ~ spl66_4), inference(avatar_split_clause, [], [f322, f614, f1003])).
fof(f322, plain, (~ (e3 = op(e3, e3)) | ~ sP60), inference(cnf_transformation, [], [f126])).
fof(f1001, plain, (~ spl66_81 | spl66_4), inference(avatar_split_clause, [], [f317, f614, f995])).
fof(f317, plain, ((e3 = op(e3, e3)) | ~ sP61), inference(cnf_transformation, [], [f125])).
fof(f125, plain, ((~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ~ sP61), inference(nnf_transformation, [], [f115])).
fof(f1000, plain, (~ spl66_81 | ~ spl66_4), inference(avatar_split_clause, [], [f318, f614, f995])).
fof(f318, plain, (~ (e3 = op(e3, e3)) | ~ sP61), inference(cnf_transformation, [], [f125])).
fof(f993, plain, (~ spl66_80 | spl66_4), inference(avatar_split_clause, [], [f313, f614, f987])).
fof(f313, plain, ((e3 = op(e3, e3)) | ~ sP62), inference(cnf_transformation, [], [f124])).
fof(f124, plain, ((~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ~ sP62), inference(nnf_transformation, [], [f116])).
fof(f992, plain, (~ spl66_80 | ~ spl66_4), inference(avatar_split_clause, [], [f314, f614, f987])).
fof(f314, plain, (~ (e3 = op(e3, e3)) | ~ sP62), inference(cnf_transformation, [], [f124])).
fof(f985, plain, (~ spl66_75 | spl66_79), inference(avatar_split_clause, [], [f305, f982, f960])).
fof(f305, plain, ((e0 = op(e0, op(e0, e0))) | ~ sP63), inference(cnf_transformation, [], [f123])).
fof(f123, plain, ((((e3 = op(e3, e0)) | ~ (e0 = op(e3, e3))) & ((e2 = op(e2, e0)) | ~ (e0 = op(e2, e2))) & ((e1 = op(e1, e0)) | ~ (e0 = op(e1, e1))) & ((e0 = op(e0, e0)) | ~ (e0 = op(e0, e0))) & (e3 = op(e0, op(e0, e3))) & (e2 = op(e0, op(e0, e2))) & (e1 = op(e0, op(e0, e1))) & (e0 = op(e0, op(e0, e0)))) | ~ sP63), inference(nnf_transformation, [], [f117])).
fof(f980, plain, (~ spl66_75 | spl66_78), inference(avatar_split_clause, [], [f306, f977, f960])).
fof(f306, plain, ((e1 = op(e0, op(e0, e1))) | ~ sP63), inference(cnf_transformation, [], [f123])).
fof(f975, plain, (~ spl66_75 | spl66_77), inference(avatar_split_clause, [], [f307, f972, f960])).
fof(f307, plain, ((e2 = op(e0, op(e0, e2))) | ~ sP63), inference(cnf_transformation, [], [f123])).
fof(f970, plain, (~ spl66_75 | spl66_76), inference(avatar_split_clause, [], [f308, f967, f960])).
fof(f308, plain, ((e3 = op(e0, op(e0, e3))) | ~ sP63), inference(cnf_transformation, [], [f123])).
fof(f965, plain, (~ spl66_75 | ~ spl66_41 | spl66_46), inference(avatar_split_clause, [], [f310, f793, f772, f960])).
fof(f310, plain, ((e1 = op(e1, e0)) | ~ (e0 = op(e1, e1)) | ~ sP63), inference(cnf_transformation, [], [f123])).
fof(f964, plain, (~ spl66_75 | ~ spl66_21 | spl66_31), inference(avatar_split_clause, [], [f311, f729, f687, f960])).
fof(f311, plain, ((e2 = op(e2, e0)) | ~ (e0 = op(e2, e2)) | ~ sP63), inference(cnf_transformation, [], [f123])).
fof(f963, plain, (~ spl66_75 | ~ spl66_1 | spl66_16), inference(avatar_split_clause, [], [f312, f665, f602, f960])).
fof(f312, plain, ((e3 = op(e3, e0)) | ~ (e0 = op(e3, e3)) | ~ sP63), inference(cnf_transformation, [], [f123])).
fof(f958, plain, (~ spl66_70 | spl66_74), inference(avatar_split_clause, [], [f297, f955, f933])).
fof(f297, plain, ((e0 = op(e1, op(e1, e0))) | ~ sP64), inference(cnf_transformation, [], [f122])).
fof(f122, plain, ((((e3 = op(e3, e1)) | ~ (e1 = op(e3, e3))) & ((e2 = op(e2, e1)) | ~ (e1 = op(e2, e2))) & ((e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & ((e0 = op(e0, e1)) | ~ (op(e0, e0) = e1)) & (e3 = op(e1, op(e1, e3))) & (e2 = op(e1, op(e1, e2))) & (e1 = op(e1, op(e1, e1))) & (e0 = op(e1, op(e1, e0)))) | ~ sP64), inference(nnf_transformation, [], [f118])).
fof(f953, plain, (~ spl66_70 | spl66_73), inference(avatar_split_clause, [], [f298, f950, f933])).
fof(f298, plain, ((e1 = op(e1, op(e1, e1))) | ~ sP64), inference(cnf_transformation, [], [f122])).
fof(f948, plain, (~ spl66_70 | spl66_72), inference(avatar_split_clause, [], [f299, f945, f933])).
fof(f299, plain, ((e2 = op(e1, op(e1, e2))) | ~ sP64), inference(cnf_transformation, [], [f122])).
fof(f943, plain, (~ spl66_70 | spl66_71), inference(avatar_split_clause, [], [f300, f940, f933])).
fof(f300, plain, ((e3 = op(e1, op(e1, e3))) | ~ sP64), inference(cnf_transformation, [], [f122])).
fof(f938, plain, (~ spl66_70 | ~ spl66_62 | spl66_57), inference(avatar_split_clause, [], [f301, f840, f861, f933])).
fof(f301, plain, ((e0 = op(e0, e1)) | ~ (op(e0, e0) = e1) | ~ sP64), inference(cnf_transformation, [], [f122])).
fof(f937, plain, (~ spl66_70 | ~ spl66_22 | spl66_27), inference(avatar_split_clause, [], [f303, f712, f691, f933])).
fof(f303, plain, ((e2 = op(e2, e1)) | ~ (e1 = op(e2, e2)) | ~ sP64), inference(cnf_transformation, [], [f122])).
fof(f936, plain, (~ spl66_70 | ~ spl66_2 | spl66_12), inference(avatar_split_clause, [], [f304, f648, f606, f933])).
fof(f304, plain, ((e3 = op(e3, e1)) | ~ (e1 = op(e3, e3)) | ~ sP64), inference(cnf_transformation, [], [f122])).
fof(f931, plain, (~ spl66_65 | spl66_69), inference(avatar_split_clause, [], [f289, f928, f906])).
fof(f289, plain, ((e0 = op(e2, op(e2, e0))) | ~ sP65), inference(cnf_transformation, [], [f121])).
fof(f121, plain, ((((e3 = op(e3, e2)) | ~ (e2 = op(e3, e3))) & ((e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & ((e1 = op(e1, e2)) | ~ (e2 = op(e1, e1))) & ((e0 = op(e0, e2)) | ~ (op(e0, e0) = e2)) & (e3 = op(e2, op(e2, e3))) & (e2 = op(e2, op(e2, e2))) & (e1 = op(e2, op(e2, e1))) & (e0 = op(e2, op(e2, e0)))) | ~ sP65), inference(nnf_transformation, [], [f119])).
fof(f926, plain, (~ spl66_65 | spl66_68), inference(avatar_split_clause, [], [f290, f923, f906])).
fof(f290, plain, ((e1 = op(e2, op(e2, e1))) | ~ sP65), inference(cnf_transformation, [], [f121])).
fof(f921, plain, (~ spl66_65 | spl66_67), inference(avatar_split_clause, [], [f291, f918, f906])).
fof(f291, plain, ((e2 = op(e2, op(e2, e2))) | ~ sP65), inference(cnf_transformation, [], [f121])).
fof(f911, plain, (~ spl66_65 | ~ spl66_63 | spl66_53), inference(avatar_split_clause, [], [f293, f823, f865, f906])).
fof(f293, plain, ((e0 = op(e0, e2)) | ~ (op(e0, e0) = e2) | ~ sP65), inference(cnf_transformation, [], [f121])).
fof(f910, plain, (~ spl66_65 | ~ spl66_43 | spl66_38), inference(avatar_split_clause, [], [f294, f759, f780, f906])).
fof(f294, plain, ((e1 = op(e1, e2)) | ~ (e2 = op(e1, e1)) | ~ sP65), inference(cnf_transformation, [], [f121])).
fof(f909, plain, (~ spl66_65 | ~ spl66_3 | spl66_8), inference(avatar_split_clause, [], [f296, f631, f610, f906])).
fof(f296, plain, ((e3 = op(e3, e2)) | ~ (e2 = op(e3, e3)) | ~ sP65), inference(cnf_transformation, [], [f121])).
fof(f904, plain, (spl66_61 | spl66_57 | spl66_53 | spl66_49), inference(avatar_split_clause, [], [f203, f806, f823, f840, f857])).
fof(f203, plain, ((e0 = op(e0, e3)) | (e0 = op(e0, e2)) | (e0 = op(e0, e1)) | (e0 = op(e0, e0))), inference(cnf_transformation, [], [f2])).
fof(f2, plain, (((e3 = op(e3, e3)) | (e3 = op(e2, e3)) | (e3 = op(e1, e3)) | (e3 = op(e0, e3))) & ((e3 = op(e3, e3)) | (e3 = op(e3, e2)) | (e3 = op(e3, e1)) | (e3 = op(e3, e0))) & ((e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(e0, e3))) & ((e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, e0))) & ((e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(e0, e3))) & ((e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, e0))) & ((e0 = op(e3, e3)) | (e0 = op(e2, e3)) | (e0 = op(e1, e3)) | (e0 = op(e0, e3))) & ((e0 = op(e3, e3)) | (e0 = op(e3, e2)) | (e0 = op(e3, e1)) | (e0 = op(e3, e0))) & ((e3 = op(e3, e2)) | (e3 = op(e2, e2)) | (e3 = op(e1, e2)) | (e3 = op(e0, e2))) & ((e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, e0))) & ((e2 = op(e3, e2)) | (e2 = op(e2, e2)) | (e2 = op(e1, e2)) | (e2 = op(e0, e2))) & ((e2 = op(e2, e3)) | (e2 = op(e2, e2)) | (e2 = op(e2, e1)) | (e2 = op(e2, e0))) & ((e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(e0, e2))) & ((e1 = op(e2, e3)) | (e1 = op(e2, e2)) | (e1 = op(e2, e1)) | (e1 = op(e2, e0))) & ((e0 = op(e3, e2)) | (e0 = op(e2, e2)) | (e0 = op(e1, e2)) | (e0 = op(e0, e2))) & ((e0 = op(e2, e3)) | (e0 = op(e2, e2)) | (e0 = op(e2, e1)) | (e0 = op(e2, e0))) & ((e3 = op(e3, e1)) | (e3 = op(e2, e1)) | (e3 = op(e1, e1)) | (e3 = op(e0, e1))) & ((e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, e0))) & ((e2 = op(e3, e1)) | (e2 = op(e2, e1)) | (e2 = op(e1, e1)) | (e2 = op(e0, e1))) & ((e2 = op(e1, e3)) | (e2 = op(e1, e2)) | (e2 = op(e1, e1)) | (e2 = op(e1, e0))) & ((e1 = op(e3, e1)) | (e1 = op(e2, e1)) | (e1 = op(e1, e1)) | (e1 = op(e0, e1))) & ((e1 = op(e1, e3)) | (e1 = op(e1, e2)) | (e1 = op(e1, e1)) | (e1 = op(e1, e0))) & ((e0 = op(e3, e1)) | (e0 = op(e2, e1)) | (e0 = op(e1, e1)) | (e0 = op(e0, e1))) & ((e0 = op(e1, e3)) | (e0 = op(e1, e2)) | (e0 = op(e1, e1)) | (e0 = op(e1, e0))) & ((e3 = op(e3, e0)) | (e3 = op(e2, e0)) | (e3 = op(e1, e0)) | (op(e0, e0) = e3)) & ((e3 = op(e0, e3)) | (e3 = op(e0, e2)) | (e3 = op(e0, e1)) | (op(e0, e0) = e3)) & ((e2 = op(e3, e0)) | (e2 = op(e2, e0)) | (e2 = op(e1, e0)) | (op(e0, e0) = e2)) & ((e2 = op(e0, e3)) | (e2 = op(e0, e2)) | (e2 = op(e0, e1)) | (op(e0, e0) = e2)) & ((e1 = op(e3, e0)) | (e1 = op(e2, e0)) | (e1 = op(e1, e0)) | (op(e0, e0) = e1)) & ((e1 = op(e0, e3)) | (e1 = op(e0, e2)) | (e1 = op(e0, e1)) | (op(e0, e0) = e1)) & ((e0 = op(e3, e0)) | (e0 = op(e2, e0)) | (e0 = op(e1, e0)) | (e0 = op(e0, e0))) & ((e0 = op(e0, e3)) | (e0 = op(e0, e2)) | (e0 = op(e0, e1)) | (e0 = op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG154+1.p', ax2)).
fof(f903, plain, (spl66_61 | spl66_45 | spl66_29 | spl66_13), inference(avatar_split_clause, [], [f204, f653, f721, f789, f857])).
fof(f204, plain, ((e0 = op(e3, e0)) | (e0 = op(e2, e0)) | (e0 = op(e1, e0)) | (e0 = op(e0, e0))), inference(cnf_transformation, [], [f2])).
fof(f902, plain, (spl66_62 | spl66_58 | spl66_54 | spl66_50), inference(avatar_split_clause, [], [f205, f810, f827, f844, f861])).
fof(f205, plain, ((e1 = op(e0, e3)) | (e1 = op(e0, e2)) | (e1 = op(e0, e1)) | (op(e0, e0) = e1)), inference(cnf_transformation, [], [f2])).
fof(f901, plain, (spl66_62 | spl66_46 | spl66_30 | spl66_14), inference(avatar_split_clause, [], [f206, f657, f725, f793, f861])).
fof(f206, plain, ((e1 = op(e3, e0)) | (e1 = op(e2, e0)) | (e1 = op(e1, e0)) | (op(e0, e0) = e1)), inference(cnf_transformation, [], [f2])).
fof(f900, plain, (spl66_63 | spl66_59 | spl66_55 | spl66_51), inference(avatar_split_clause, [], [f207, f814, f831, f848, f865])).
fof(f207, plain, ((e2 = op(e0, e3)) | (e2 = op(e0, e2)) | (e2 = op(e0, e1)) | (op(e0, e0) = e2)), inference(cnf_transformation, [], [f2])).
fof(f899, plain, (spl66_63 | spl66_47 | spl66_31 | spl66_15), inference(avatar_split_clause, [], [f208, f661, f729, f797, f865])).
fof(f208, plain, ((e2 = op(e3, e0)) | (e2 = op(e2, e0)) | (e2 = op(e1, e0)) | (op(e0, e0) = e2)), inference(cnf_transformation, [], [f2])).
fof(f898, plain, (spl66_64 | spl66_60 | spl66_56 | spl66_52), inference(avatar_split_clause, [], [f209, f818, f835, f852, f869])).
fof(f209, plain, ((e3 = op(e0, e3)) | (e3 = op(e0, e2)) | (e3 = op(e0, e1)) | (op(e0, e0) = e3)), inference(cnf_transformation, [], [f2])).
fof(f897, plain, (spl66_64 | spl66_48 | spl66_32 | spl66_16), inference(avatar_split_clause, [], [f210, f665, f733, f801, f869])).
fof(f210, plain, ((e3 = op(e3, e0)) | (e3 = op(e2, e0)) | (e3 = op(e1, e0)) | (op(e0, e0) = e3)), inference(cnf_transformation, [], [f2])).
fof(f896, plain, (spl66_45 | spl66_41 | spl66_37 | spl66_33), inference(avatar_split_clause, [], [f211, f738, f755, f772, f789])).
fof(f211, plain, ((e0 = op(e1, e3)) | (e0 = op(e1, e2)) | (e0 = op(e1, e1)) | (e0 = op(e1, e0))), inference(cnf_transformation, [], [f2])).
fof(f895, plain, (spl66_57 | spl66_41 | spl66_25 | spl66_9), inference(avatar_split_clause, [], [f212, f636, f704, f772, f840])).
fof(f212, plain, ((e0 = op(e3, e1)) | (e0 = op(e2, e1)) | (e0 = op(e1, e1)) | (e0 = op(e0, e1))), inference(cnf_transformation, [], [f2])).
fof(f894, plain, (spl66_46 | spl66_42 | spl66_38 | spl66_34), inference(avatar_split_clause, [], [f213, f742, f759, f776, f793])).
fof(f213, plain, ((e1 = op(e1, e3)) | (e1 = op(e1, e2)) | (e1 = op(e1, e1)) | (e1 = op(e1, e0))), inference(cnf_transformation, [], [f2])).
fof(f892, plain, (spl66_47 | spl66_43 | spl66_39 | spl66_35), inference(avatar_split_clause, [], [f215, f746, f763, f780, f797])).
fof(f215, plain, ((e2 = op(e1, e3)) | (e2 = op(e1, e2)) | (e2 = op(e1, e1)) | (e2 = op(e1, e0))), inference(cnf_transformation, [], [f2])).
fof(f891, plain, (spl66_59 | spl66_43 | spl66_27 | spl66_11), inference(avatar_split_clause, [], [f216, f644, f712, f780, f848])).
fof(f216, plain, ((e2 = op(e3, e1)) | (e2 = op(e2, e1)) | (e2 = op(e1, e1)) | (e2 = op(e0, e1))), inference(cnf_transformation, [], [f2])).
fof(f890, plain, (spl66_48 | spl66_44 | spl66_40 | spl66_36), inference(avatar_split_clause, [], [f217, f750, f767, f784, f801])).
fof(f217, plain, ((e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, e0))), inference(cnf_transformation, [], [f2])).
fof(f889, plain, (spl66_60 | spl66_44 | spl66_28 | spl66_12), inference(avatar_split_clause, [], [f218, f648, f716, f784, f852])).
fof(f218, plain, ((e3 = op(e3, e1)) | (e3 = op(e2, e1)) | (e3 = op(e1, e1)) | (e3 = op(e0, e1))), inference(cnf_transformation, [], [f2])).
fof(f888, plain, (spl66_29 | spl66_25 | spl66_21 | spl66_17), inference(avatar_split_clause, [], [f219, f670, f687, f704, f721])).
fof(f219, plain, ((e0 = op(e2, e3)) | (e0 = op(e2, e2)) | (e0 = op(e2, e1)) | (e0 = op(e2, e0))), inference(cnf_transformation, [], [f2])).
fof(f887, plain, (spl66_53 | spl66_37 | spl66_21 | spl66_5), inference(avatar_split_clause, [], [f220, f619, f687, f755, f823])).
fof(f220, plain, ((e0 = op(e3, e2)) | (e0 = op(e2, e2)) | (e0 = op(e1, e2)) | (e0 = op(e0, e2))), inference(cnf_transformation, [], [f2])).
fof(f886, plain, (spl66_30 | spl66_26 | spl66_22 | spl66_18), inference(avatar_split_clause, [], [f221, f674, f691, f708, f725])).
fof(f221, plain, ((e1 = op(e2, e3)) | (e1 = op(e2, e2)) | (e1 = op(e2, e1)) | (e1 = op(e2, e0))), inference(cnf_transformation, [], [f2])).
fof(f885, plain, (spl66_54 | spl66_38 | spl66_22 | spl66_6), inference(avatar_split_clause, [], [f222, f623, f691, f759, f827])).
fof(f222, plain, ((e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(e0, e2))), inference(cnf_transformation, [], [f2])).
fof(f884, plain, (spl66_31 | spl66_27 | spl66_23 | spl66_19), inference(avatar_split_clause, [], [f223, f678, f695, f712, f729])).
fof(f223, plain, ((e2 = op(e2, e3)) | (e2 = op(e2, e2)) | (e2 = op(e2, e1)) | (e2 = op(e2, e0))), inference(cnf_transformation, [], [f2])).
fof(f882, plain, (spl66_32 | spl66_28 | spl66_24 | spl66_20), inference(avatar_split_clause, [], [f225, f682, f699, f716, f733])).
fof(f225, plain, ((e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, e0))), inference(cnf_transformation, [], [f2])).
fof(f881, plain, (spl66_56 | spl66_40 | spl66_24 | spl66_8), inference(avatar_split_clause, [], [f226, f631, f699, f767, f835])).
fof(f226, plain, ((e3 = op(e3, e2)) | (e3 = op(e2, e2)) | (e3 = op(e1, e2)) | (e3 = op(e0, e2))), inference(cnf_transformation, [], [f2])).
fof(f880, plain, (spl66_13 | spl66_9 | spl66_5 | spl66_1), inference(avatar_split_clause, [], [f227, f602, f619, f636, f653])).
fof(f227, plain, ((e0 = op(e3, e3)) | (e0 = op(e3, e2)) | (e0 = op(e3, e1)) | (e0 = op(e3, e0))), inference(cnf_transformation, [], [f2])).
fof(f879, plain, (spl66_49 | spl66_33 | spl66_17 | spl66_1), inference(avatar_split_clause, [], [f228, f602, f670, f738, f806])).
fof(f228, plain, ((e0 = op(e3, e3)) | (e0 = op(e2, e3)) | (e0 = op(e1, e3)) | (e0 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f878, plain, (spl66_14 | spl66_10 | spl66_6 | spl66_2), inference(avatar_split_clause, [], [f229, f606, f623, f640, f657])).
fof(f229, plain, ((e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, e0))), inference(cnf_transformation, [], [f2])).
fof(f877, plain, (spl66_50 | spl66_34 | spl66_18 | spl66_2), inference(avatar_split_clause, [], [f230, f606, f674, f742, f810])).
fof(f230, plain, ((e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f876, plain, (spl66_15 | spl66_11 | spl66_7 | spl66_3), inference(avatar_split_clause, [], [f231, f610, f627, f644, f661])).
fof(f231, plain, ((e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, e0))), inference(cnf_transformation, [], [f2])).
fof(f875, plain, (spl66_51 | spl66_35 | spl66_19 | spl66_3), inference(avatar_split_clause, [], [f232, f610, f678, f746, f814])).
fof(f232, plain, ((e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f874, plain, (spl66_16 | spl66_12 | spl66_8 | spl66_4), inference(avatar_split_clause, [], [f233, f614, f631, f648, f665])).
fof(f233, plain, ((e3 = op(e3, e3)) | (e3 = op(e3, e2)) | (e3 = op(e3, e1)) | (e3 = op(e3, e0))), inference(cnf_transformation, [], [f2])).
fof(f873, plain, (spl66_52 | spl66_36 | spl66_20 | spl66_4), inference(avatar_split_clause, [], [f234, f614, f682, f750, f818])).
fof(f234, plain, ((e3 = op(e3, e3)) | (e3 = op(e2, e3)) | (e3 = op(e1, e3)) | (e3 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f872, plain, (spl66_61 | spl66_62 | spl66_63 | spl66_64), inference(avatar_split_clause, [], [f187, f869, f865, f861, f857])).
fof(f187, plain, ((op(e0, e0) = e3) | (op(e0, e0) = e2) | (op(e0, e0) = e1) | (e0 = op(e0, e0))), inference(cnf_transformation, [], [f1])).
fof(f1, plain, (((e3 = op(e3, e3)) | (e2 = op(e3, e3)) | (e1 = op(e3, e3)) | (e0 = op(e3, e3))) & ((e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (e0 = op(e3, e2))) & ((e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (e0 = op(e3, e1))) & ((e3 = op(e3, e0)) | (e2 = op(e3, e0)) | (e1 = op(e3, e0)) | (e0 = op(e3, e0))) & ((e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))) & ((e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (e0 = op(e2, e2))) & ((e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))) & ((e3 = op(e2, e0)) | (e2 = op(e2, e0)) | (e1 = op(e2, e0)) | (e0 = op(e2, e0))) & ((e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (e0 = op(e1, e3))) & ((e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (e0 = op(e1, e2))) & ((e3 = op(e1, e1)) | (e2 = op(e1, e1)) | (e1 = op(e1, e1)) | (e0 = op(e1, e1))) & ((e3 = op(e1, e0)) | (e2 = op(e1, e0)) | (e1 = op(e1, e0)) | (e0 = op(e1, e0))) & ((e3 = op(e0, e3)) | (e2 = op(e0, e3)) | (e1 = op(e0, e3)) | (e0 = op(e0, e3))) & ((e3 = op(e0, e2)) | (e2 = op(e0, e2)) | (e1 = op(e0, e2)) | (e0 = op(e0, e2))) & ((e3 = op(e0, e1)) | (e2 = op(e0, e1)) | (e1 = op(e0, e1)) | (e0 = op(e0, e1))) & ((op(e0, e0) = e3) | (op(e0, e0) = e2) | (op(e0, e0) = e1) | (e0 = op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG154+1.p', ax1)).
fof(f855, plain, (spl66_57 | spl66_58 | spl66_59 | spl66_60), inference(avatar_split_clause, [], [f188, f852, f848, f844, f840])).
fof(f188, plain, ((e3 = op(e0, e1)) | (e2 = op(e0, e1)) | (e1 = op(e0, e1)) | (e0 = op(e0, e1))), inference(cnf_transformation, [], [f1])).
fof(f838, plain, (spl66_53 | spl66_54 | spl66_55 | spl66_56), inference(avatar_split_clause, [], [f189, f835, f831, f827, f823])).
fof(f189, plain, ((e3 = op(e0, e2)) | (e2 = op(e0, e2)) | (e1 = op(e0, e2)) | (e0 = op(e0, e2))), inference(cnf_transformation, [], [f1])).
fof(f821, plain, (spl66_49 | spl66_50 | spl66_51 | spl66_52), inference(avatar_split_clause, [], [f190, f818, f814, f810, f806])).
fof(f190, plain, ((e3 = op(e0, e3)) | (e2 = op(e0, e3)) | (e1 = op(e0, e3)) | (e0 = op(e0, e3))), inference(cnf_transformation, [], [f1])).
fof(f804, plain, (spl66_45 | spl66_46 | spl66_47 | spl66_48), inference(avatar_split_clause, [], [f191, f801, f797, f793, f789])).
fof(f191, plain, ((e3 = op(e1, e0)) | (e2 = op(e1, e0)) | (e1 = op(e1, e0)) | (e0 = op(e1, e0))), inference(cnf_transformation, [], [f1])).
fof(f787, plain, (spl66_41 | spl66_42 | spl66_43 | spl66_44), inference(avatar_split_clause, [], [f192, f784, f780, f776, f772])).
fof(f192, plain, ((e3 = op(e1, e1)) | (e2 = op(e1, e1)) | (e1 = op(e1, e1)) | (e0 = op(e1, e1))), inference(cnf_transformation, [], [f1])).
fof(f770, plain, (spl66_37 | spl66_38 | spl66_39 | spl66_40), inference(avatar_split_clause, [], [f193, f767, f763, f759, f755])).
fof(f193, plain, ((e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (e0 = op(e1, e2))), inference(cnf_transformation, [], [f1])).
fof(f753, plain, (spl66_33 | spl66_34 | spl66_35 | spl66_36), inference(avatar_split_clause, [], [f194, f750, f746, f742, f738])).
fof(f194, plain, ((e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (e0 = op(e1, e3))), inference(cnf_transformation, [], [f1])).
fof(f736, plain, (spl66_29 | spl66_30 | spl66_31 | spl66_32), inference(avatar_split_clause, [], [f195, f733, f729, f725, f721])).
fof(f195, plain, ((e3 = op(e2, e0)) | (e2 = op(e2, e0)) | (e1 = op(e2, e0)) | (e0 = op(e2, e0))), inference(cnf_transformation, [], [f1])).
fof(f719, plain, (spl66_25 | spl66_26 | spl66_27 | spl66_28), inference(avatar_split_clause, [], [f196, f716, f712, f708, f704])).
fof(f196, plain, ((e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))), inference(cnf_transformation, [], [f1])).
fof(f702, plain, (spl66_21 | spl66_22 | spl66_23 | spl66_24), inference(avatar_split_clause, [], [f197, f699, f695, f691, f687])).
fof(f197, plain, ((e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (e0 = op(e2, e2))), inference(cnf_transformation, [], [f1])).
fof(f685, plain, (spl66_17 | spl66_18 | spl66_19 | spl66_20), inference(avatar_split_clause, [], [f198, f682, f678, f674, f670])).
fof(f198, plain, ((e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))), inference(cnf_transformation, [], [f1])).
fof(f668, plain, (spl66_13 | spl66_14 | spl66_15 | spl66_16), inference(avatar_split_clause, [], [f199, f665, f661, f657, f653])).
fof(f199, plain, ((e3 = op(e3, e0)) | (e2 = op(e3, e0)) | (e1 = op(e3, e0)) | (e0 = op(e3, e0))), inference(cnf_transformation, [], [f1])).
fof(f651, plain, (spl66_9 | spl66_10 | spl66_11 | spl66_12), inference(avatar_split_clause, [], [f200, f648, f644, f640, f636])).
fof(f200, plain, ((e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (e0 = op(e3, e1))), inference(cnf_transformation, [], [f1])).
fof(f634, plain, (spl66_5 | spl66_6 | spl66_7 | spl66_8), inference(avatar_split_clause, [], [f201, f631, f627, f623, f619])).
fof(f201, plain, ((e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (e0 = op(e3, e2))), inference(cnf_transformation, [], [f1])).
fof(f617, plain, (spl66_1 | spl66_2 | spl66_3 | spl66_4), inference(avatar_split_clause, [], [f202, f614, f610, f606, f602])).
fof(f202, plain, ((e3 = op(e3, e3)) | (e2 = op(e3, e3)) | (e1 = op(e3, e3)) | (e0 = op(e3, e3))), inference(cnf_transformation, [], [f1])).