fof(f12781, plain, $false, inference(avatar_sat_refutation, [], [f631, f652, f673, f694, f715, f736, f757, f778, f799, f820, f841, f862, f883, f904, f925, f946, f967, f988, f1009, f1030, f1051, f1072, f1093, f1114, f1135, f1136, f1137, f1138, f1139, f1140, f1141, f1142, f1143, f1145, f1146, f1147, f1148, f1149, f1150, f1151, f1152, f1153, f1154, f1155, f1156, f1157, f1158, f1159, f1161, f1162, f1163, f1164, f1166, f1167, f1168, f1169, f1170, f1171, f1172, f1173, f1174, f1175, f1176, f1177, f1178, f1179, f1180, f1181, f1182, f1183, f1184, f1185, f1191, f1192, f1199, f1201, f1202, f1203, f1210, f1211, f1212, f1217, f1220, f1221, f1222, f1223, f1224, f1226, f1252, f1270, f1280, f1290, f1313, f1328, f1361, f1362, f1371, f1394, f1417, f1422, f1435, f1441, f1442, f1447, f1469, f1472, f1473, f1474, f1484, f1487, f1488, f1489, f1490, f1495, f1501, f1503, f1511, f1516, f1522, f1523, f1526, f1535, f1538, f1539, f1541, f1546, f1604, f1607, f1629, f1643, f1653, f1670, f1676, f1692, f1714, f1757, f1760, f1763, f1810, f1830, f1840, f1843, f1846, f1871, f1877, f1900, f1915, f1927, f1930, f1936, f1943, f1946, f1949, f1952, f1960, f1963, f1967, f1973, f1977, f1980, f1983, f1993, f2011, f2014, f2020, f2035, f2038, f2051, f2059, f2066, f2067, f2069, f2093, f2109, f2120, f2130, f2134, f2137, f2146, f2160, f2168, f2170, f2188, f2189, f2218, f2230, f2248, f2259, f2273, f2276, f2290, f2307, f2321, f2323, f2361, f2369, f2378, f2397, f2398, f2436, f2458, f2482, f2498, f2509, f2519, f2564, f2582, f2585, f2605, f2616, f2634, f2636, f2656, f2658, f2676, f2711, f2716, f2719, f2734, f2754, f2759, f2769, f2782, f2790, f2793, f2800, f2808, f2834, f2842, f2871, f2886, f2888, f2898, f2899, f2903, f2914, f2933, f2946, f2952, f2995, f2998, f3000, f3014, f3015, f3018, f3023, f3029, f3041, f3042, f3043, f3064, f3068, f3083, f3098, f3110, f3114, f3168, f3182, f3184, f3205, f3227, f3248, f3258, f3265, f3267, f3272, f3281, f3293, f3311, f3345, f3355, f3359, f3382, f3386, f3411, f3441, f3503, f3509, f3515, f3519, f3545, f3548, f3565, f3566, f3576, f3579, f3591, f3593, f3601, f3609, f3613, f3629, f3636, f3639, f3642, f3649, f3651, f3660, f3666, f3677, f3683, f3701, f3711, f3732, f3734, f3737, f3741, f3761, f3786, f3804, f3820, f3830, f3832, f3834, f3849, f3904, f3913, f3914, f3944, f3976, f3985, f4009, f4035, f4040, f4053, f4064, f4075, f4100, f4121, f4133, f4141, f4170, f4177, f4193, f4194, f4212, f4231, f4245, f4283, f4305, f4306, f4308, f4338, f4347, f4353, f4357, f4360, f4373, f4382, f4388, f4401, f4405, f4407, f4430, f4460, f4479, f4481, f4499, f4522, f4531, f4556, f4565, f4567, f4575, f4576, f4579, f4589, f4608, f4620, f4633, f4635, f4637, f4642, f4644, f4652, f4669, f4687, f4721, f4738, f4740, f4758, f4762, f4768, f4771, f4787, f4799, f4871, f4876, f4885, f4888, f4936, f4955, f4967, f4970, f4972, f4980, f5005, f5007, f5036, f5048, f5085, f5092, f5094, f5117, f5142, f5144, f5147, f5156, f5158, f5160, f5168, f5172, f5223, f5234, f5236, f5237, f5238, f5269, f5282, f5293, f5296, f5329, f5356, f5365, f5367, f5382, f5405, f5416, f5428, f5432, f5466, f5488, f5528, f5539, f5540, f5605, f5606, f5621, f5651, f5659, f5668, f5675, f5678, f5680, f5686, f5714, f5727, f5741, f5749, f5750, f5764, f5769, f5789, f5803, f5828, f5835, f5836, f5840, f5847, f5861, f5867, f5881, f5883, f5892, f5895, f5912, f5928, f5950, f5973, f5992, f6001, f6033, f6034, f6070, f6071, f6085, f6091, f6094, f6114, f6120, f6139, f6187, f6190, f6193, f6197, f6207, f6214, f6279, f6292, f6294, f6304, f6315, f6316, f6338, f6352, f6395, f6407, f6409, f6410, f6431, f6438, f6456, f6463, f6469, f6472, f6482, f6526, f6540, f6554, f6571, f6578, f6584, f6619, f6643, f6647, f6678, f6704, f6713, f6716, f6720, f6748, f6769, f6787, f6798, f6800, f6841, f6842, f6863, f6875, f6900, f6902, f6907, f6912, f6913, f6953, f6954, f6970, f6991, f7008, f7014, f7017, f7058, f7073, f7085, f7094, f7117, f7119, f7129, f7178, f7188, f7203, f7216, f7217, f7228, f7283, f7314, f7336, f7346, f7350, f7358, f7368, f7376, f7395, f7409, f7430, f7438, f7470, f7473, f7474, f7487, f7502, f7514, f7534, f7545, f7583, f7593, f7599, f7606, f7611, f7612, f7619, f7629, f7634, f7635, f7650, f7671, f7692, f7695, f7697, f7714, f7718, f7737, f7747, f7751, f7754, f7798, f7822, f7834, f7836, f7846, f7901, f7932, f7955, f7984, f8017, f8032, f8051, f8077, f8083, f8119, f8160, f8175, f8195, f8198, f8241, f8266, f8267, f8284, f8285, f8306, f8307, f8316, f8321, f8331, f8333, f8334, f8341, f8361, f8394, f8412, f8413, f8432, f8434, f8437, f8464, f8473, f8475, f8531, f8532, f8535, f8536, f8544, f8552, f8556, f8580, f8581, f8585, f8596, f8598, f8614, f8615, f8633, f8647, f8677, f8718, f8818, f8820, f8832, f8841, f8865, f8881, f8890, f8900, f8907, f8973, f8975, f8976, f9003, f9010, f9027, f9028, f9029, f9030, f9048, f9049, f9055, f9057, f9061, f9092, f9093, f9107, f9120, f9140, f9147, f9162, f9169, f9209, f9257, f9260, f9263, f9274, f9281, f9298, f9336, f9338, f9348, f9357, f9358, f9365, f9385, f9441, f9459, f9486, f9495, f9516, f9517, f9518, f9536, f9550, f9551, f9560, f9575, f9588, f9589, f9603, f9646, f9662, f9674, f9677, f9696, f9708, f9709, f9717, f9718, f9728, f9741, f9750, f9751, f9760, f9761, f9763, f9779, f9780, f9790, f9791, f9803, f9817, f9830, f9862, f9867, f9886, f9890, f9906, f9922, f9932, f9953, f9983, f9993, f10018, f10019, f10030, f10045, f10052, f10065, f10075, f10103, f10112, f10122, f10132, f10134, f10136, f10138, f10140, f10187, f10188, f10199, f10200, f10211, f10228, f10240, f10241, f10242, f10278, f10279, f10299, f10300, f10316, f10321, f10328, f10374, f10384, f10392, f10403, f10417, f10426, f10441, f10442, f10443, f10450, f10451, f10463, f10464, f10473, f10483, f10484, f10485, f10494, f10504, f10511, f10512, f10520, f10535, f10536, f10544, f10545, f10546, f10573, f10576, f10579, f10599, f10626, f10639, f10641, f10665, f10673, f10693, f10694, f10695, f10697, f10718, f10740, f10741, f10742, f10755, f10756, f10764, f10765, f10774, f10775, f10785, f10808, f10809, f10810, f10823, f10824, f10839, f10840, f10845, f10859, f10874, f10883, f10928, f10929, f10944, f10954, f10969, f10970, f10980, f10990, f11009, f11011, f11012, f11027, f11029, f11039, f11040, f11047, f11061, f11070, f11072, f11087, f11094, f11103, f11104, f11113, f11124, f11131, f11132, f11139, f11140, f11149, f11151, f11176, f11187, f11198, f11202, f11235, f11299, f11301, f11337, f11338, f11348, f11349, f11359, f11365, f11378, f11379, f11380, f11397, f11398, f11399, f11426, f11434, f11443, f11452, f11462, f11470, f11476, f11482, f11485, f11489, f11494, f11504, f11513, f11542, f11543, f11565, f11576, f11583, f11590, f11599, f11600, f11601, f11611, f11612, f11621, f11632, f11640, f11641, f11650, f11666, f11671, f11679, f11680, f11706, f11720, f11763, f11783, f11784, f11785, f11786, f11800, f11820, f11827, f11828, f11837, f11838, f11839, f11848, f11849, f11850, f11855, f11872, f11874, f11883, f11897, f11907, f11908, f11916, f11960, f11961, f11967, f11985, f11989, f11996, f12002, f12039, f12069, f12070, f12082, f12093, f12113, f12114, f12125, f12126, f12138, f12139, f12140, f12147, f12152, f12153, f12154, f12165, f12174, f12175, f12194, f12195, f12210, f12211, f12212, f12222, f12239, f12240, f12253, f12255, f12257, f12267, f12286, f12323, f12341, f12342, f12343, f12353, f12354, f12366, f12367, f12375, f12384, f12397, f12408, f12409, f12422, f12423, f12451, f12452, f12460, f12482, f12490, f12496, f12500, f12519, f12528, f12545, f12547, f12562, f12563, f12573, f12581, f12590, f12600, f12601, f12602, f12619, f12620, f12621, f12649, f12650, f12660, f12671, f12672, f12678, f12679, f12690, f12704, f12705, f12706, f12712, f12723, f12746, f12761])).
fof(f12761, plain, (~ spl4_3 | ~ spl4_13), inference(avatar_split_clause, [], [f12756, f662, f620])).
fof(f620, plain, (spl4_3 <=> (e2 = op(e4, e4))), introduced(avatar_definition, [new_symbols(naming, [spl4_3])])).
fof(f662, plain, (spl4_13 <=> (e2 = op(e4, e2))), introduced(avatar_definition, [new_symbols(naming, [spl4_13])])).
fof(f12756, plain, (~ (e2 = op(e4, e4)) | ~ spl4_13), inference(backward_demodulation, [], [f454, f664])).
fof(f664, plain, ((e2 = op(e4, e2)) | ~ spl4_13), inference(avatar_component_clause, [], [f662])).
fof(f454, plain, ~ (op(e4, e2) = op(e4, e4)), inference(cnf_transformation, [], [f4])).
fof(f4, plain, (~ (op(e4, e3) = op(e4, e4)) & ~ (op(e4, e2) = op(e4, e4)) & ~ (op(e4, e2) = op(e4, e3)) & ~ (op(e4, e1) = op(e4, e4)) & ~ (op(e4, e1) = op(e4, e3)) & ~ (op(e4, e1) = op(e4, e2)) & ~ (op(e4, e0) = op(e4, e4)) & ~ (op(e4, e0) = op(e4, e3)) & ~ (op(e4, e0) = op(e4, e2)) & ~ (op(e4, e0) = op(e4, e1)) & ~ (op(e3, e3) = op(e3, e4)) & ~ (op(e3, e2) = op(e3, e4)) & ~ (op(e3, e2) = op(e3, e3)) & ~ (op(e3, e1) = op(e3, e4)) & ~ (op(e3, e1) = op(e3, e3)) & ~ (op(e3, e1) = op(e3, e2)) & ~ (op(e3, e0) = op(e3, e4)) & ~ (op(e3, e0) = op(e3, e3)) & ~ (op(e3, e0) = op(e3, e2)) & ~ (op(e3, e0) = op(e3, e1)) & ~ (op(e2, e3) = op(e2, e4)) & ~ (op(e2, e2) = op(e2, e4)) & ~ (op(e2, e2) = op(e2, e3)) & ~ (op(e2, e1) = op(e2, e4)) & ~ (op(e2, e1) = op(e2, e3)) & ~ (op(e2, e1) = op(e2, e2)) & ~ (op(e2, e0) = op(e2, e4)) & ~ (op(e2, e0) = op(e2, e3)) & ~ (op(e2, e0) = op(e2, e2)) & ~ (op(e2, e0) = op(e2, e1)) & ~ (op(e1, e3) = op(e1, e4)) & ~ (op(e1, e2) = op(e1, e4)) & ~ (op(e1, e2) = op(e1, e3)) & ~ (op(e1, e1) = op(e1, e4)) & ~ (op(e1, e1) = op(e1, e3)) & ~ (op(e1, e1) = op(e1, e2)) & ~ (op(e1, e0) = op(e1, e4)) & ~ (op(e1, e0) = op(e1, e3)) & ~ (op(e1, e0) = op(e1, e2)) & ~ (op(e1, e0) = op(e1, e1)) & ~ (op(e0, e3) = op(e0, e4)) & ~ (op(e0, e2) = op(e0, e4)) & ~ (op(e0, e2) = op(e0, e3)) & ~ (op(e0, e1) = op(e0, e4)) & ~ (op(e0, e1) = op(e0, e3)) & ~ (op(e0, e1) = op(e0, e2)) & ~ (op(e0, e0) = op(e0, e4)) & ~ (op(e0, e0) = op(e0, e3)) & ~ (op(e0, e0) = op(e0, e2)) & ~ (op(e0, e0) = op(e0, e1)) & ~ (op(e3, e4) = op(e4, e4)) & ~ (op(e2, e4) = op(e4, e4)) & ~ (op(e2, e4) = op(e3, e4)) & ~ (op(e1, e4) = op(e4, e4)) & ~ (op(e1, e4) = op(e3, e4)) & ~ (op(e1, e4) = op(e2, e4)) & ~ (op(e0, e4) = op(e4, e4)) & ~ (op(e0, e4) = op(e3, e4)) & ~ (op(e0, e4) = op(e2, e4)) & ~ (op(e0, e4) = op(e1, e4)) & ~ (op(e3, e3) = op(e4, e3)) & ~ (op(e2, e3) = op(e4, e3)) & ~ (op(e2, e3) = op(e3, e3)) & ~ (op(e1, e3) = op(e4, e3)) & ~ (op(e1, e3) = op(e3, e3)) & ~ (op(e1, e3) = op(e2, e3)) & ~ (op(e0, e3) = op(e4, e3)) & ~ (op(e0, e3) = op(e3, e3)) & ~ (op(e0, e3) = op(e2, e3)) & ~ (op(e0, e3) = op(e1, e3)) & ~ (op(e3, e2) = op(e4, e2)) & ~ (op(e2, e2) = op(e4, e2)) & ~ (op(e2, e2) = op(e3, e2)) & ~ (op(e1, e2) = op(e4, e2)) & ~ (op(e1, e2) = op(e3, e2)) & ~ (op(e1, e2) = op(e2, e2)) & ~ (op(e0, e2) = op(e4, e2)) & ~ (op(e0, e2) = op(e3, e2)) & ~ (op(e0, e2) = op(e2, e2)) & ~ (op(e0, e2) = op(e1, e2)) & ~ (op(e3, e1) = op(e4, e1)) & ~ (op(e2, e1) = op(e4, e1)) & ~ (op(e2, e1) = op(e3, e1)) & ~ (op(e1, e1) = op(e4, e1)) & ~ (op(e1, e1) = op(e3, e1)) & ~ (op(e1, e1) = op(e2, e1)) & ~ (op(e0, e1) = op(e4, e1)) & ~ (op(e0, e1) = op(e3, e1)) & ~ (op(e0, e1) = op(e2, e1)) & ~ (op(e0, e1) = op(e1, e1)) & ~ (op(e3, e0) = op(e4, e0)) & ~ (op(e2, e0) = op(e4, e0)) & ~ (op(e2, e0) = op(e3, e0)) & ~ (op(e1, e0) = op(e4, e0)) & ~ (op(e1, e0) = op(e3, e0)) & ~ (op(e1, e0) = op(e2, e0)) & ~ (op(e0, e0) = op(e4, e0)) & ~ (op(e0, e0) = op(e3, e0)) & ~ (op(e0, e0) = op(e2, e0)) & ~ (op(e0, e0) = op(e1, e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG191+1.p', ax4)).
fof(f12746, plain, (~ spl4_11 | ~ spl4_21), inference(avatar_split_clause, [], [f12737, f696, f654])).
fof(f654, plain, (spl4_11 <=> (e0 = op(e4, e2))), introduced(avatar_definition, [new_symbols(naming, [spl4_11])])).
fof(f696, plain, (spl4_21 <=> (e0 = op(e4, e0))), introduced(avatar_definition, [new_symbols(naming, [spl4_21])])).
fof(f12737, plain, (~ (e0 = op(e4, e2)) | ~ spl4_21), inference(backward_demodulation, [], [f447, f698])).
fof(f698, plain, ((e0 = op(e4, e0)) | ~ spl4_21), inference(avatar_component_clause, [], [f696])).
fof(f447, plain, ~ (op(e4, e0) = op(e4, e2)), inference(cnf_transformation, [], [f4])).
fof(f12723, plain, (~ spl4_27 | ~ spl4_32), inference(avatar_split_clause, [], [f12715, f742, f721])).
fof(f721, plain, (spl4_27 <=> (e1 = op(e3, e4))), introduced(avatar_definition, [new_symbols(naming, [spl4_27])])).
fof(f742, plain, (spl4_32 <=> (e1 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl4_32])])).
fof(f12715, plain, (~ (e1 = op(e3, e4)) | ~ spl4_32), inference(backward_demodulation, [], [f445, f744])).
fof(f744, plain, ((e1 = op(e3, e3)) | ~ spl4_32), inference(avatar_component_clause, [], [f742])).
fof(f445, plain, ~ (op(e3, e3) = op(e3, e4)), inference(cnf_transformation, [], [f4])).
fof(f12712, plain, (spl4_13 | ~ spl4_40 | ~ spl4_58), inference(avatar_split_clause, [], [f12710, f851, f775, f662])).
fof(f775, plain, (spl4_40 <=> (e4 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl4_40])])).
fof(f851, plain, (spl4_58 <=> (e2 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl4_58])])).
fof(f12710, plain, ((e2 = op(e4, e2)) | (~ spl4_40 | ~ spl4_58)), inference(backward_demodulation, [], [f12691, f777])).
fof(f777, plain, ((e4 = op(e3, e2)) | ~ spl4_40), inference(avatar_component_clause, [], [f775])).
fof(f12691, plain, ((e2 = op(op(e3, e2), e2)) | ~ spl4_58), inference(backward_demodulation, [], [f344, f853])).
fof(f853, plain, ((e2 = op(e2, e3)) | ~ spl4_58), inference(avatar_component_clause, [], [f851])).
fof(f344, plain, (e2 = op(op(e3, e2), op(e2, e3))), inference(cnf_transformation, [], [f3])).
fof(f3, plain, ((e4 = op(op(e4, e4), op(e4, e4))) & (e4 = op(op(e3, e4), op(e4, e3))) & (e4 = op(op(e2, e4), op(e4, e2))) & (e4 = op(op(e1, e4), op(e4, e1))) & (e4 = op(op(e0, e4), op(e4, e0))) & (e3 = op(op(e4, e3), op(e3, e4))) & (e3 = op(op(e3, e3), op(e3, e3))) & (e3 = op(op(e2, e3), op(e3, e2))) & (e3 = op(op(e1, e3), op(e3, e1))) & (e3 = op(op(e0, e3), op(e3, e0))) & (e2 = op(op(e4, e2), op(e2, e4))) & (e2 = op(op(e3, e2), op(e2, e3))) & (e2 = op(op(e2, e2), op(e2, e2))) & (e2 = op(op(e1, e2), op(e2, e1))) & (e2 = op(op(e0, e2), op(e2, e0))) & (e1 = op(op(e4, e1), op(e1, e4))) & (e1 = op(op(e3, e1), op(e1, e3))) & (e1 = op(op(e2, e1), op(e1, e2))) & (e1 = op(op(e1, e1), op(e1, e1))) & (e1 = op(op(e0, e1), op(e1, e0))) & (e0 = op(op(e4, e0), op(e0, e4))) & (e0 = op(op(e3, e0), op(e0, e3))) & (e0 = op(op(e2, e0), op(e0, e2))) & (e0 = op(op(e1, e0), op(e0, e1))) & (e0 = op(op(e0, e0), op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG191+1.p', ax3)).
fof(f12706, plain, (spl4_9 | ~ spl4_49 | ~ spl4_110), inference(avatar_split_clause, [], [f12703, f1069, f813, f645])).
fof(f645, plain, (spl4_9 <=> (e3 = op(e4, e3))), introduced(avatar_definition, [new_symbols(naming, [spl4_9])])).
fof(f813, plain, (spl4_49 <=> (e3 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl4_49])])).
fof(f1069, plain, (spl4_110 <=> (e4 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl4_110])])).
fof(f12703, plain, ((e3 = op(e4, e3)) | (~ spl4_49 | ~ spl4_110)), inference(backward_demodulation, [], [f12575, f815])).
fof(f815, plain, ((e3 = op(e3, e0)) | ~ spl4_49), inference(avatar_component_clause, [], [f813])).
fof(f12575, plain, ((e3 = op(e4, op(e3, e0))) | ~ spl4_110), inference(backward_demodulation, [], [f346, f1071])).
fof(f1071, plain, ((e4 = op(e0, e3)) | ~ spl4_110), inference(avatar_component_clause, [], [f1069])).
fof(f346, plain, (e3 = op(op(e0, e3), op(e3, e0))), inference(cnf_transformation, [], [f3])).
fof(f12705, plain, (spl4_26 | ~ spl4_49 | ~ spl4_110), inference(avatar_split_clause, [], [f12702, f1069, f813, f717])).
fof(f717, plain, (spl4_26 <=> (e0 = op(e3, e4))), introduced(avatar_definition, [new_symbols(naming, [spl4_26])])).
fof(f12702, plain, ((e0 = op(e3, e4)) | (~ spl4_49 | ~ spl4_110)), inference(backward_demodulation, [], [f12574, f815])).
fof(f12574, plain, ((e0 = op(op(e3, e0), e4)) | ~ spl4_110), inference(backward_demodulation, [], [f334, f1071])).
fof(f334, plain, (e0 = op(op(e3, e0), op(e0, e3))), inference(cnf_transformation, [], [f3])).
fof(f12704, plain, (~ spl4_39 | ~ spl4_49), inference(avatar_split_clause, [], [f12698, f813, f771])).
fof(f771, plain, (spl4_39 <=> (e3 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl4_39])])).
fof(f12698, plain, (~ (e3 = op(e3, e2)) | ~ spl4_49), inference(backward_demodulation, [], [f437, f815])).
fof(f437, plain, ~ (op(e3, e0) = op(e3, e2)), inference(cnf_transformation, [], [f4])).
fof(f12690, plain, (~ spl4_11 | ~ spl4_61), inference(avatar_split_clause, [], [f12682, f864, f654])).
fof(f864, plain, (spl4_61 <=> (e0 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl4_61])])).
fof(f12682, plain, (~ (e0 = op(e4, e2)) | ~ spl4_61), inference(backward_demodulation, [], [f384, f866])).
fof(f866, plain, ((e0 = op(e2, e2)) | ~ spl4_61), inference(avatar_component_clause, [], [f864])).
fof(f384, plain, ~ (op(e2, e2) = op(e4, e2)), inference(cnf_transformation, [], [f4])).
fof(f12679, plain, (spl4_17 | ~ spl4_70 | ~ spl4_87), inference(avatar_split_clause, [], [f12676, f973, f901, f679])).
fof(f679, plain, (spl4_17 <=> (e1 = op(e4, e1))), introduced(avatar_definition, [new_symbols(naming, [spl4_17])])).
fof(f901, plain, (spl4_70 <=> (e4 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl4_70])])).
fof(f973, plain, (spl4_87 <=> (e1 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl4_87])])).
fof(f12676, plain, ((e1 = op(e4, e1)) | (~ spl4_70 | ~ spl4_87)), inference(backward_demodulation, [], [f12624, f903])).
fof(f903, plain, ((e4 = op(e2, e1)) | ~ spl4_70), inference(avatar_component_clause, [], [f901])).
fof(f12624, plain, ((e1 = op(op(e2, e1), e1)) | ~ spl4_87), inference(backward_demodulation, [], [f338, f975])).
fof(f975, plain, ((e1 = op(e1, e2)) | ~ spl4_87), inference(avatar_component_clause, [], [f973])).
fof(f338, plain, (e1 = op(op(e2, e1), op(e1, e2))), inference(cnf_transformation, [], [f3])).
fof(f12678, plain, (~ spl4_65 | ~ spl4_70), inference(avatar_split_clause, [], [f12674, f901, f880])).
fof(f880, plain, (spl4_65 <=> (e4 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl4_65])])).
fof(f12674, plain, (~ (e4 = op(e2, e2)) | ~ spl4_70), inference(backward_demodulation, [], [f430, f903])).
fof(f430, plain, ~ (op(e2, e1) = op(e2, e2)), inference(cnf_transformation, [], [f4])).
fof(f12672, plain, (~ spl4_57 | ~ spl4_72), inference(avatar_split_clause, [], [f12665, f910, f847])).
fof(f847, plain, (spl4_57 <=> (e1 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl4_57])])).
fof(f910, plain, (spl4_72 <=> (e1 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl4_72])])).
fof(f12665, plain, (~ (e1 = op(e2, e3)) | ~ spl4_72), inference(backward_demodulation, [], [f428, f912])).
fof(f912, plain, ((e1 = op(e2, e0)) | ~ spl4_72), inference(avatar_component_clause, [], [f910])).
fof(f428, plain, ~ (op(e2, e0) = op(e2, e3)), inference(cnf_transformation, [], [f4])).
fof(f12671, plain, (~ spl4_22 | ~ spl4_72), inference(avatar_split_clause, [], [f12662, f910, f700])).
fof(f700, plain, (spl4_22 <=> (e1 = op(e4, e0))), introduced(avatar_definition, [new_symbols(naming, [spl4_22])])).
fof(f12662, plain, (~ (e1 = op(e4, e0)) | ~ spl4_72), inference(backward_demodulation, [], [f364, f912])).
fof(f364, plain, ~ (op(e2, e0) = op(e4, e0)), inference(cnf_transformation, [], [f4])).
fof(f12660, plain, (~ spl4_3 | ~ spl4_78), inference(avatar_split_clause, [], [f12657, f935, f620])).
fof(f935, plain, (spl4_78 <=> (e2 = op(e1, e4))), introduced(avatar_definition, [new_symbols(naming, [spl4_78])])).
fof(f12657, plain, (~ (e2 = op(e4, e4)) | ~ spl4_78), inference(backward_demodulation, [], [f402, f937])).
fof(f937, plain, ((e2 = op(e1, e4)) | ~ spl4_78), inference(avatar_component_clause, [], [f935])).
fof(f402, plain, ~ (op(e1, e4) = op(e4, e4)), inference(cnf_transformation, [], [f4])).
fof(f12650, plain, (spl4_72 | ~ spl4_43 | ~ spl4_81), inference(avatar_split_clause, [], [f12643, f948, f788, f910])).
fof(f788, plain, (spl4_43 <=> (e2 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl4_43])])).
fof(f948, plain, (spl4_81 <=> (e0 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl4_81])])).
fof(f12643, plain, ((e1 = op(e2, e0)) | (~ spl4_43 | ~ spl4_81)), inference(backward_demodulation, [], [f12509, f950])).
fof(f950, plain, ((e0 = op(e1, e3)) | ~ spl4_81), inference(avatar_component_clause, [], [f948])).
fof(f12509, plain, ((e1 = op(e2, op(e1, e3))) | ~ spl4_43), inference(forward_demodulation, [], [f339, f790])).
fof(f790, plain, ((e2 = op(e3, e1)) | ~ spl4_43), inference(avatar_component_clause, [], [f788])).
fof(f339, plain, (e1 = op(op(e3, e1), op(e1, e3))), inference(cnf_transformation, [], [f3])).
fof(f12649, plain, (~ spl4_31 | ~ spl4_81), inference(avatar_split_clause, [], [f12635, f948, f738])).
fof(f738, plain, (spl4_31 <=> (e0 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl4_31])])).
fof(f12635, plain, (~ (e0 = op(e3, e3)) | ~ spl4_81), inference(backward_demodulation, [], [f391, f950])).
fof(f391, plain, ~ (op(e1, e3) = op(e3, e3)), inference(cnf_transformation, [], [f4])).
fof(f12621, plain, (~ spl4_84 | ~ spl4_94), inference(avatar_split_clause, [], [f12607, f1002, f960])).
fof(f960, plain, (spl4_84 <=> (e3 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl4_84])])).
fof(f1002, plain, (spl4_94 <=> (e3 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl4_94])])).
fof(f12607, plain, (~ (e3 = op(e1, e3)) | ~ spl4_94), inference(backward_demodulation, [], [f421, f1004])).
fof(f1004, plain, ((e3 = op(e1, e1)) | ~ spl4_94), inference(avatar_component_clause, [], [f1002])).
fof(f421, plain, ~ (op(e1, e1) = op(e1, e3)), inference(cnf_transformation, [], [f4])).
fof(f12620, plain, (~ spl4_19 | ~ spl4_94), inference(avatar_split_clause, [], [f12605, f1002, f687])).
fof(f687, plain, (spl4_19 <=> (e3 = op(e4, e1))), introduced(avatar_definition, [new_symbols(naming, [spl4_19])])).
fof(f12605, plain, (~ (e3 = op(e4, e1)) | ~ spl4_94), inference(backward_demodulation, [], [f372, f1004])).
fof(f372, plain, ~ (op(e1, e1) = op(e4, e1)), inference(cnf_transformation, [], [f4])).
fof(f12619, plain, (spl4_32 | ~ spl4_94), inference(avatar_split_clause, [], [f12603, f1002, f742])).
fof(f12603, plain, ((e1 = op(e3, e3)) | ~ spl4_94), inference(backward_demodulation, [], [f337, f1004])).
fof(f337, plain, (e1 = op(op(e1, e1), op(e1, e1))), inference(cnf_transformation, [], [f3])).
fof(f12602, plain, (spl4_21 | ~ spl4_100 | ~ spl4_116), inference(avatar_split_clause, [], [f12598, f1095, f1027, f696])).
fof(f1027, plain, (spl4_100 <=> (e4 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl4_100])])).
fof(f1095, plain, (spl4_116 <=> (e0 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl4_116])])).
fof(f12598, plain, ((e0 = op(e4, e0)) | (~ spl4_100 | ~ spl4_116)), inference(backward_demodulation, [], [f12550, f1029])).
fof(f1029, plain, ((e4 = op(e1, e0)) | ~ spl4_100), inference(avatar_component_clause, [], [f1027])).
fof(f12550, plain, ((e0 = op(op(e1, e0), e0)) | ~ spl4_116), inference(backward_demodulation, [], [f332, f1097])).
fof(f1097, plain, ((e0 = op(e0, e1)) | ~ spl4_116), inference(avatar_component_clause, [], [f1095])).
fof(f332, plain, (e0 = op(op(e1, e0), op(e0, e1))), inference(cnf_transformation, [], [f3])).
fof(f12601, plain, (~ spl4_80 | ~ spl4_100), inference(avatar_split_clause, [], [f12597, f1027, f943])).
fof(f943, plain, (spl4_80 <=> (e4 = op(e1, e4))), introduced(avatar_definition, [new_symbols(naming, [spl4_80])])).
fof(f12597, plain, (~ (e4 = op(e1, e4)) | ~ spl4_100), inference(backward_demodulation, [], [f419, f1029])).
fof(f419, plain, ~ (op(e1, e0) = op(e1, e4)), inference(cnf_transformation, [], [f4])).
fof(f12600, plain, (~ spl4_50 | ~ spl4_100), inference(avatar_split_clause, [], [f12592, f1027, f817])).
fof(f817, plain, (spl4_50 <=> (e4 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl4_50])])).
fof(f12592, plain, (~ (e4 = op(e3, e0)) | ~ spl4_100), inference(backward_demodulation, [], [f361, f1029])).
fof(f361, plain, ~ (op(e1, e0) = op(e3, e0)), inference(cnf_transformation, [], [f4])).
fof(f12590, plain, (~ spl4_27 | ~ spl4_102), inference(avatar_split_clause, [], [f12585, f1036, f721])).
fof(f1036, plain, (spl4_102 <=> (e1 = op(e0, e4))), introduced(avatar_definition, [new_symbols(naming, [spl4_102])])).
fof(f12585, plain, (~ (e1 = op(e3, e4)) | ~ spl4_102), inference(backward_demodulation, [], [f398, f1038])).
fof(f1038, plain, ((e1 = op(e0, e4)) | ~ spl4_102), inference(avatar_component_clause, [], [f1036])).
fof(f398, plain, ~ (op(e0, e4) = op(e3, e4)), inference(cnf_transformation, [], [f4])).
fof(f12581, plain, (~ spl4_10 | ~ spl4_110), inference(avatar_split_clause, [], [f12579, f1069, f649])).
fof(f649, plain, (spl4_10 <=> (e4 = op(e4, e3))), introduced(avatar_definition, [new_symbols(naming, [spl4_10])])).
fof(f12579, plain, (~ (e4 = op(e4, e3)) | ~ spl4_110), inference(backward_demodulation, [], [f389, f1071])).
fof(f389, plain, ~ (op(e0, e3) = op(e4, e3)), inference(cnf_transformation, [], [f4])).
fof(f12573, plain, (~ spl4_39 | ~ spl4_114), inference(avatar_split_clause, [], [f12568, f1086, f771])).
fof(f1086, plain, (spl4_114 <=> (e3 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl4_114])])).
fof(f12568, plain, (~ (e3 = op(e3, e2)) | ~ spl4_114), inference(backward_demodulation, [], [f378, f1088])).
fof(f1088, plain, ((e3 = op(e0, e2)) | ~ spl4_114), inference(avatar_component_clause, [], [f1086])).
fof(f378, plain, ~ (op(e0, e2) = op(e3, e2)), inference(cnf_transformation, [], [f4])).
fof(f12563, plain, (~ spl4_101 | ~ spl4_116), inference(avatar_split_clause, [], [f12557, f1095, f1032])).
fof(f1032, plain, (spl4_101 <=> (e0 = op(e0, e4))), introduced(avatar_definition, [new_symbols(naming, [spl4_101])])).
fof(f12557, plain, (~ (e0 = op(e0, e4)) | ~ spl4_116), inference(backward_demodulation, [], [f412, f1097])).
fof(f412, plain, ~ (op(e0, e1) = op(e0, e4)), inference(cnf_transformation, [], [f4])).
fof(f12562, plain, (~ spl4_66 | ~ spl4_116), inference(avatar_split_clause, [], [f12553, f1095, f885])).
fof(f885, plain, (spl4_66 <=> (e0 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl4_66])])).
fof(f12553, plain, (~ (e0 = op(e2, e1)) | ~ spl4_116), inference(backward_demodulation, [], [f367, f1097])).
fof(f367, plain, ~ (op(e0, e1) = op(e2, e1)), inference(cnf_transformation, [], [f4])).
fof(f12547, plain, (~ spl4_108 | ~ spl4_123), inference(avatar_split_clause, [], [f12538, f1124, f1061])).
fof(f1061, plain, (spl4_108 <=> (e2 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl4_108])])).
fof(f1124, plain, (spl4_123 <=> (op(e0, e0) = e2)), introduced(avatar_definition, [new_symbols(naming, [spl4_123])])).
fof(f12538, plain, (~ (e2 = op(e0, e3)) | ~ spl4_123), inference(backward_demodulation, [], [f408, f1126])).
fof(f1126, plain, ((op(e0, e0) = e2) | ~ spl4_123), inference(avatar_component_clause, [], [f1124])).
fof(f408, plain, ~ (op(e0, e0) = op(e0, e3)), inference(cnf_transformation, [], [f4])).
fof(f12545, plain, (spl4_61 | ~ spl4_123), inference(avatar_split_clause, [], [f12531, f1124, f864])).
fof(f12531, plain, ((e0 = op(e2, e2)) | ~ spl4_123), inference(backward_demodulation, [], [f331, f1126])).
fof(f331, plain, (e0 = op(op(e0, e0), op(e0, e0))), inference(cnf_transformation, [], [f3])).
fof(f12528, plain, (~ spl4_70 | ~ spl4_154 | spl4_155), inference(avatar_split_clause, [], [f12526, f1346, f1341, f901])).
fof(f1341, plain, (spl4_154 <=> (e1 = op(e2, op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl4_154])])).
fof(f1346, plain, (spl4_155 <=> (e4 = op(e2, op(e2, op(e2, e2))))), introduced(avatar_definition, [new_symbols(naming, [spl4_155])])).
fof(f12526, plain, (~ (e4 = op(e2, e1)) | (~ spl4_154 | spl4_155)), inference(backward_demodulation, [], [f1348, f1342])).
fof(f1342, plain, ((e1 = op(e2, op(e2, e2))) | ~ spl4_154), inference(avatar_component_clause, [], [f1341])).
fof(f1348, plain, (~ (e4 = op(e2, op(e2, op(e2, e2)))) | spl4_155), inference(avatar_component_clause, [], [f1346])).
fof(f12519, plain, (~ spl4_110 | ~ spl4_168 | spl4_173), inference(avatar_split_clause, [], [f12518, f1444, f1414, f1069])).
fof(f1414, plain, (spl4_168 <=> (e3 = op(e0, op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl4_168])])).
fof(f1444, plain, (spl4_173 <=> (e4 = op(e0, op(e0, op(e0, e0))))), introduced(avatar_definition, [new_symbols(naming, [spl4_173])])).
fof(f12518, plain, (~ (e4 = op(e0, e3)) | (~ spl4_168 | spl4_173)), inference(forward_demodulation, [], [f1446, f1415])).
fof(f1415, plain, ((e3 = op(e0, op(e0, e0))) | ~ spl4_168), inference(avatar_component_clause, [], [f1414])).
fof(f1446, plain, (~ (e4 = op(e0, op(e0, op(e0, e0)))) | spl4_173), inference(avatar_component_clause, [], [f1444])).
fof(f12500, plain, (~ spl4_28 | ~ spl4_43), inference(avatar_split_clause, [], [f12499, f788, f725])).
fof(f725, plain, (spl4_28 <=> (e2 = op(e3, e4))), introduced(avatar_definition, [new_symbols(naming, [spl4_28])])).
fof(f12499, plain, (~ (e2 = op(e3, e4)) | ~ spl4_43), inference(forward_demodulation, [], [f442, f790])).
fof(f442, plain, ~ (op(e3, e1) = op(e3, e4)), inference(cnf_transformation, [], [f4])).
fof(f12496, plain, (~ spl4_27 | spl4_131 | ~ spl4_132), inference(avatar_split_clause, [], [f12327, f1236, f1232, f721])).
fof(f1232, plain, (spl4_131 <=> (e1 = op(e3, op(e3, op(e3, e3))))), introduced(avatar_definition, [new_symbols(naming, [spl4_131])])).
fof(f1236, plain, (spl4_132 <=> (e4 = op(e3, op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl4_132])])).
fof(f12327, plain, (~ (e1 = op(e3, e4)) | (spl4_131 | ~ spl4_132)), inference(forward_demodulation, [], [f1234, f1237])).
fof(f1237, plain, ((e4 = op(e3, op(e3, e3))) | ~ spl4_132), inference(avatar_component_clause, [], [f1236])).
fof(f1234, plain, (~ (e1 = op(e3, op(e3, op(e3, e3)))) | spl4_131), inference(avatar_component_clause, [], [f1232])).
fof(f12490, plain, (~ spl4_1 | ~ spl4_11), inference(avatar_split_clause, [], [f12483, f654, f612])).
fof(f612, plain, (spl4_1 <=> (e0 = op(e4, e4))), introduced(avatar_definition, [new_symbols(naming, [spl4_1])])).
fof(f12483, plain, (~ (e0 = op(e4, e4)) | ~ spl4_11), inference(backward_demodulation, [], [f454, f656])).
fof(f656, plain, ((e0 = op(e4, e2)) | ~ spl4_11), inference(avatar_component_clause, [], [f654])).
fof(f12482, plain, (~ spl4_12 | ~ spl4_22), inference(avatar_split_clause, [], [f12475, f700, f658])).
fof(f658, plain, (spl4_12 <=> (e1 = op(e4, e2))), introduced(avatar_definition, [new_symbols(naming, [spl4_12])])).
fof(f12475, plain, (~ (e1 = op(e4, e2)) | ~ spl4_22), inference(backward_demodulation, [], [f447, f702])).
fof(f702, plain, ((e1 = op(e4, e0)) | ~ spl4_22), inference(avatar_component_clause, [], [f700])).
fof(f12460, plain, (spl4_39 | ~ spl4_43 | ~ spl4_84), inference(avatar_split_clause, [], [f12458, f960, f788, f771])).
fof(f12458, plain, ((e3 = op(e3, e2)) | (~ spl4_43 | ~ spl4_84)), inference(backward_demodulation, [], [f12142, f790])).
fof(f12142, plain, ((e3 = op(e3, op(e3, e1))) | ~ spl4_84), inference(backward_demodulation, [], [f347, f962])).
fof(f962, plain, ((e3 = op(e1, e3)) | ~ spl4_84), inference(avatar_component_clause, [], [f960])).
fof(f347, plain, (e3 = op(op(e1, e3), op(e3, e1))), inference(cnf_transformation, [], [f3])).
fof(f12452, plain, (spl4_11 | ~ spl4_50 | ~ spl4_108), inference(avatar_split_clause, [], [f12449, f1061, f817, f654])).
fof(f12449, plain, ((e0 = op(e4, e2)) | (~ spl4_50 | ~ spl4_108)), inference(backward_demodulation, [], [f12368, f819])).
fof(f819, plain, ((e4 = op(e3, e0)) | ~ spl4_50), inference(avatar_component_clause, [], [f817])).
fof(f12368, plain, ((e0 = op(op(e3, e0), e2)) | ~ spl4_108), inference(backward_demodulation, [], [f334, f1063])).
fof(f1063, plain, ((e2 = op(e0, e3)) | ~ spl4_108), inference(avatar_component_clause, [], [f1061])).
fof(f12451, plain, (~ spl4_40 | ~ spl4_50), inference(avatar_split_clause, [], [f12447, f817, f775])).
fof(f12447, plain, (~ (e4 = op(e3, e2)) | ~ spl4_50), inference(backward_demodulation, [], [f437, f819])).
fof(f12423, plain, (~ spl4_40 | ~ spl4_65), inference(avatar_split_clause, [], [f12411, f880, f775])).
fof(f12411, plain, (~ (e4 = op(e3, e2)) | ~ spl4_65), inference(backward_demodulation, [], [f383, f882])).
fof(f882, plain, ((e4 = op(e2, e2)) | ~ spl4_65), inference(avatar_component_clause, [], [f880])).
fof(f383, plain, ~ (op(e2, e2) = op(e3, e2)), inference(cnf_transformation, [], [f4])).
fof(f12422, plain, (spl4_3 | ~ spl4_65), inference(avatar_split_clause, [], [f12410, f880, f620])).
fof(f12410, plain, ((e2 = op(e4, e4)) | ~ spl4_65), inference(backward_demodulation, [], [f343, f882])).
fof(f343, plain, (e2 = op(op(e2, e2), op(e2, e2))), inference(cnf_transformation, [], [f3])).
fof(f12409, plain, (~ spl4_56 | ~ spl4_66), inference(avatar_split_clause, [], [f12400, f885, f843])).
fof(f843, plain, (spl4_56 <=> (e0 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl4_56])])).
fof(f12400, plain, (~ (e0 = op(e2, e3)) | ~ spl4_66), inference(backward_demodulation, [], [f431, f887])).
fof(f887, plain, ((e0 = op(e2, e1)) | ~ spl4_66), inference(avatar_component_clause, [], [f885])).
fof(f431, plain, ~ (op(e2, e1) = op(e2, e3)), inference(cnf_transformation, [], [f4])).
fof(f12408, plain, (~ spl4_41 | ~ spl4_66), inference(avatar_split_clause, [], [f12398, f885, f780])).
fof(f780, plain, (spl4_41 <=> (e0 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl4_41])])).
fof(f12398, plain, (~ (e0 = op(e3, e1)) | ~ spl4_66), inference(backward_demodulation, [], [f373, f887])).
fof(f373, plain, ~ (op(e2, e1) = op(e3, e1)), inference(cnf_transformation, [], [f4])).
fof(f12397, plain, (spl4_66 | ~ spl4_73 | ~ spl4_112), inference(avatar_split_clause, [], [f12393, f1078, f914, f885])).
fof(f914, plain, (spl4_73 <=> (e2 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl4_73])])).
fof(f1078, plain, (spl4_112 <=> (e1 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl4_112])])).
fof(f12393, plain, ((e0 = op(e2, e1)) | (~ spl4_73 | ~ spl4_112)), inference(backward_demodulation, [], [f12356, f916])).
fof(f916, plain, ((e2 = op(e2, e0)) | ~ spl4_73), inference(avatar_component_clause, [], [f914])).
fof(f12356, plain, ((e0 = op(op(e2, e0), e1)) | ~ spl4_112), inference(backward_demodulation, [], [f333, f1080])).
fof(f1080, plain, ((e1 = op(e0, e2)) | ~ spl4_112), inference(avatar_component_clause, [], [f1078])).
fof(f333, plain, (e0 = op(op(e2, e0), op(e0, e2))), inference(cnf_transformation, [], [f3])).
fof(f12384, plain, (~ spl4_1 | ~ spl4_101), inference(avatar_split_clause, [], [f12379, f1032, f612])).
fof(f12379, plain, (~ (e0 = op(e4, e4)) | ~ spl4_101), inference(backward_demodulation, [], [f399, f1034])).
fof(f1034, plain, ((e0 = op(e0, e4)) | ~ spl4_101), inference(avatar_component_clause, [], [f1032])).
fof(f399, plain, ~ (op(e0, e4) = op(e4, e4)), inference(cnf_transformation, [], [f4])).
fof(f12375, plain, (~ spl4_33 | ~ spl4_108), inference(avatar_split_clause, [], [f12371, f1061, f746])).
fof(f746, plain, (spl4_33 <=> (e2 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl4_33])])).
fof(f12371, plain, (~ (e2 = op(e3, e3)) | ~ spl4_108), inference(backward_demodulation, [], [f388, f1063])).
fof(f388, plain, ~ (op(e0, e3) = op(e3, e3)), inference(cnf_transformation, [], [f4])).
fof(f12367, plain, (~ spl4_107 | ~ spl4_112), inference(avatar_split_clause, [], [f12361, f1078, f1057])).
fof(f1057, plain, (spl4_107 <=> (e1 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl4_107])])).
fof(f12361, plain, (~ (e1 = op(e0, e3)) | ~ spl4_112), inference(backward_demodulation, [], [f413, f1080])).
fof(f413, plain, ~ (op(e0, e2) = op(e0, e3)), inference(cnf_transformation, [], [f4])).
fof(f12366, plain, (~ spl4_12 | ~ spl4_112), inference(avatar_split_clause, [], [f12360, f1078, f658])).
fof(f12360, plain, (~ (e1 = op(e4, e2)) | ~ spl4_112), inference(backward_demodulation, [], [f379, f1080])).
fof(f379, plain, ~ (op(e0, e2) = op(e4, e2)), inference(cnf_transformation, [], [f4])).
fof(f12354, plain, (spl4_22 | ~ spl4_96 | ~ spl4_120), inference(avatar_split_clause, [], [f12351, f1111, f1011, f700])).
fof(f1011, plain, (spl4_96 <=> (e0 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl4_96])])).
fof(f1111, plain, (spl4_120 <=> (e4 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl4_120])])).
fof(f12351, plain, ((e1 = op(e4, e0)) | (~ spl4_96 | ~ spl4_120)), inference(backward_demodulation, [], [f12279, f1113])).
fof(f1113, plain, ((e4 = op(e0, e1)) | ~ spl4_120), inference(avatar_component_clause, [], [f1111])).
fof(f12279, plain, ((e1 = op(op(e0, e1), e0)) | ~ spl4_96), inference(forward_demodulation, [], [f336, f1013])).
fof(f1013, plain, ((e0 = op(e1, e0)) | ~ spl4_96), inference(avatar_component_clause, [], [f1011])).
fof(f336, plain, (e1 = op(op(e0, e1), op(e1, e0))), inference(cnf_transformation, [], [f3])).
fof(f12353, plain, (~ spl4_70 | ~ spl4_120), inference(avatar_split_clause, [], [f12346, f1111, f901])).
fof(f12346, plain, (~ (e4 = op(e2, e1)) | ~ spl4_120), inference(backward_demodulation, [], [f367, f1113])).
fof(f12343, plain, (~ spl4_104 | ~ spl4_124), inference(avatar_split_clause, [], [f12337, f1128, f1044])).
fof(f1044, plain, (spl4_104 <=> (e3 = op(e0, e4))), introduced(avatar_definition, [new_symbols(naming, [spl4_104])])).
fof(f1128, plain, (spl4_124 <=> (op(e0, e0) = e3)), introduced(avatar_definition, [new_symbols(naming, [spl4_124])])).
fof(f12337, plain, (~ (e3 = op(e0, e4)) | ~ spl4_124), inference(backward_demodulation, [], [f409, f1130])).
fof(f1130, plain, ((op(e0, e0) = e3) | ~ spl4_124), inference(avatar_component_clause, [], [f1128])).
fof(f409, plain, ~ (op(e0, e0) = op(e0, e4)), inference(cnf_transformation, [], [f4])).
fof(f12342, plain, (~ spl4_49 | ~ spl4_124), inference(avatar_split_clause, [], [f12332, f1128, f813])).
fof(f12332, plain, (~ (e3 = op(e3, e0)) | ~ spl4_124), inference(backward_demodulation, [], [f358, f1130])).
fof(f358, plain, ~ (op(e0, e0) = op(e3, e0)), inference(cnf_transformation, [], [f4])).
fof(f12341, plain, (spl4_31 | ~ spl4_124), inference(avatar_split_clause, [], [f12330, f1128, f738])).
fof(f12330, plain, ((e0 = op(e3, e3)) | ~ spl4_124), inference(backward_demodulation, [], [f331, f1130])).
fof(f12323, plain, (~ spl4_23 | spl4_143 | ~ spl4_167), inference(avatar_split_clause, [], [f12320, f1409, f1287, f704])).
fof(f704, plain, (spl4_23 <=> (e2 = op(e4, e0))), introduced(avatar_definition, [new_symbols(naming, [spl4_23])])).
fof(f1287, plain, (spl4_143 <=> (e2 = op(e4, op(e4, op(e4, e4))))), introduced(avatar_definition, [new_symbols(naming, [spl4_143])])).
fof(f1409, plain, (spl4_167 <=> (e0 = op(e4, op(e4, e4)))), introduced(avatar_definition, [new_symbols(naming, [spl4_167])])).
fof(f12320, plain, (~ (e2 = op(e4, e0)) | (spl4_143 | ~ spl4_167)), inference(backward_demodulation, [], [f1289, f1410])).
fof(f1410, plain, ((e0 = op(e4, op(e4, e4))) | ~ spl4_167), inference(avatar_component_clause, [], [f1409])).
fof(f1289, plain, (~ (e2 = op(e4, op(e4, op(e4, e4)))) | spl4_143), inference(avatar_component_clause, [], [f1287])).
fof(f12286, plain, (~ spl4_108 | spl4_165 | ~ spl4_168), inference(avatar_split_clause, [], [f12285, f1414, f1400, f1061])).
fof(f1400, plain, (spl4_165 <=> (e2 = op(e0, op(e0, op(e0, e0))))), introduced(avatar_definition, [new_symbols(naming, [spl4_165])])).
fof(f12285, plain, (~ (e2 = op(e0, e3)) | (spl4_165 | ~ spl4_168)), inference(forward_demodulation, [], [f1402, f1415])).
fof(f1402, plain, (~ (e2 = op(e0, op(e0, op(e0, e0)))) | spl4_165), inference(avatar_component_clause, [], [f1400])).
fof(f12267, plain, (~ spl4_59 | ~ spl4_84), inference(avatar_split_clause, [], [f12266, f960, f855])).
fof(f855, plain, (spl4_59 <=> (e3 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl4_59])])).
fof(f12266, plain, (~ (e3 = op(e2, e3)) | ~ spl4_84), inference(forward_demodulation, [], [f390, f962])).
fof(f390, plain, ~ (op(e1, e3) = op(e2, e3)), inference(cnf_transformation, [], [f4])).
fof(f12257, plain, (~ spl4_5 | ~ spl4_10), inference(avatar_split_clause, [], [f12256, f649, f628])).
fof(f628, plain, (spl4_5 <=> (e4 = op(e4, e4))), introduced(avatar_definition, [new_symbols(naming, [spl4_5])])).
fof(f12256, plain, (~ (e4 = op(e4, e4)) | ~ spl4_10), inference(forward_demodulation, [], [f455, f651])).
fof(f651, plain, ((e4 = op(e4, e3)) | ~ spl4_10), inference(avatar_component_clause, [], [f649])).
fof(f455, plain, ~ (op(e4, e3) = op(e4, e4)), inference(cnf_transformation, [], [f4])).
fof(f12255, plain, (~ spl4_2 | ~ spl4_27), inference(avatar_split_clause, [], [f12254, f721, f616])).
fof(f616, plain, (spl4_2 <=> (e1 = op(e4, e4))), introduced(avatar_definition, [new_symbols(naming, [spl4_2])])).
fof(f12254, plain, (~ (e1 = op(e4, e4)) | ~ spl4_27), inference(forward_demodulation, [], [f405, f723])).
fof(f723, plain, ((e1 = op(e3, e4)) | ~ spl4_27), inference(avatar_component_clause, [], [f721])).
fof(f405, plain, ~ (op(e3, e4) = op(e4, e4)), inference(cnf_transformation, [], [f4])).
fof(f12253, plain, (~ spl4_5 | ~ spl4_80), inference(avatar_split_clause, [], [f12252, f943, f628])).
fof(f12252, plain, (~ (e4 = op(e4, e4)) | ~ spl4_80), inference(forward_demodulation, [], [f402, f945])).
fof(f945, plain, ((e4 = op(e1, e4)) | ~ spl4_80), inference(avatar_component_clause, [], [f943])).
fof(f12240, plain, (~ spl4_12 | ~ spl4_1 | ~ spl4_23 | spl4_134), inference(avatar_split_clause, [], [f12236, f1245, f704, f612, f658])).
fof(f1245, plain, (spl4_134 <=> (e1 = op(e4, op(e4, op(e4, e4))))), introduced(avatar_definition, [new_symbols(naming, [spl4_134])])).
fof(f12236, plain, (~ (e1 = op(e4, e2)) | (~ spl4_1 | ~ spl4_23 | spl4_134)), inference(backward_demodulation, [], [f12049, f706])).
fof(f706, plain, ((e2 = op(e4, e0)) | ~ spl4_23), inference(avatar_component_clause, [], [f704])).
fof(f12049, plain, (~ (e1 = op(e4, op(e4, e0))) | (~ spl4_1 | spl4_134)), inference(forward_demodulation, [], [f1247, f614])).
fof(f614, plain, ((e0 = op(e4, e4)) | ~ spl4_1), inference(avatar_component_clause, [], [f612])).
fof(f1247, plain, (~ (e1 = op(e4, op(e4, op(e4, e4)))) | spl4_134), inference(avatar_component_clause, [], [f1245])).
fof(f12239, plain, (~ spl4_8 | ~ spl4_23), inference(avatar_split_clause, [], [f12229, f704, f641])).
fof(f641, plain, (spl4_8 <=> (e2 = op(e4, e3))), introduced(avatar_definition, [new_symbols(naming, [spl4_8])])).
fof(f12229, plain, (~ (e2 = op(e4, e3)) | ~ spl4_23), inference(backward_demodulation, [], [f448, f706])).
fof(f448, plain, ~ (op(e4, e0) = op(e4, e3)), inference(cnf_transformation, [], [f4])).
fof(f12222, plain, (~ spl4_8 | ~ spl4_33), inference(avatar_split_clause, [], [f12214, f746, f641])).
fof(f12214, plain, (~ (e2 = op(e4, e3)) | ~ spl4_33), inference(backward_demodulation, [], [f395, f748])).
fof(f748, plain, ((e2 = op(e3, e3)) | ~ spl4_33), inference(avatar_component_clause, [], [f746])).
fof(f395, plain, ~ (op(e3, e3) = op(e4, e3)), inference(cnf_transformation, [], [f4])).
fof(f12212, plain, (spl4_23 | ~ spl4_40 | ~ spl4_56), inference(avatar_split_clause, [], [f12208, f843, f775, f704])).
fof(f12208, plain, ((e2 = op(e4, e0)) | (~ spl4_40 | ~ spl4_56)), inference(backward_demodulation, [], [f11730, f777])).
fof(f11730, plain, ((e2 = op(op(e3, e2), e0)) | ~ spl4_56), inference(forward_demodulation, [], [f344, f845])).
fof(f845, plain, ((e0 = op(e2, e3)) | ~ spl4_56), inference(avatar_component_clause, [], [f843])).
fof(f12211, plain, (~ spl4_30 | ~ spl4_40), inference(avatar_split_clause, [], [f12207, f775, f733])).
fof(f733, plain, (spl4_30 <=> (e4 = op(e3, e4))), introduced(avatar_definition, [new_symbols(naming, [spl4_30])])).
fof(f12207, plain, (~ (e4 = op(e3, e4)) | ~ spl4_40), inference(backward_demodulation, [], [f444, f777])).
fof(f444, plain, ~ (op(e3, e2) = op(e3, e4)), inference(cnf_transformation, [], [f4])).
fof(f12210, plain, (~ spl4_15 | ~ spl4_40), inference(avatar_split_clause, [], [f12205, f775, f670])).
fof(f670, plain, (spl4_15 <=> (e4 = op(e4, e2))), introduced(avatar_definition, [new_symbols(naming, [spl4_15])])).
fof(f12205, plain, (~ (e4 = op(e4, e2)) | ~ spl4_40), inference(backward_demodulation, [], [f385, f777])).
fof(f385, plain, ~ (op(e3, e2) = op(e4, e2)), inference(cnf_transformation, [], [f4])).
fof(f12195, plain, (spl4_41 | ~ spl4_49 | ~ spl4_107), inference(avatar_split_clause, [], [f12192, f1057, f813, f780])).
fof(f12192, plain, ((e0 = op(e3, e1)) | (~ spl4_49 | ~ spl4_107)), inference(backward_demodulation, [], [f11811, f815])).
fof(f11811, plain, ((e0 = op(op(e3, e0), e1)) | ~ spl4_107), inference(backward_demodulation, [], [f334, f1059])).
fof(f1059, plain, ((e1 = op(e0, e3)) | ~ spl4_107), inference(avatar_component_clause, [], [f1057])).
fof(f12194, plain, (~ spl4_34 | ~ spl4_49), inference(avatar_split_clause, [], [f12189, f813, f750])).
fof(f750, plain, (spl4_34 <=> (e3 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl4_34])])).
fof(f12189, plain, (~ (e3 = op(e3, e3)) | ~ spl4_49), inference(backward_demodulation, [], [f438, f815])).
fof(f438, plain, ~ (op(e3, e0) = op(e3, e3)), inference(cnf_transformation, [], [f4])).
fof(f12175, plain, (~ spl4_54 | ~ spl4_64), inference(avatar_split_clause, [], [f12169, f876, f834])).
fof(f834, plain, (spl4_54 <=> (e3 = op(e2, e4))), introduced(avatar_definition, [new_symbols(naming, [spl4_54])])).
fof(f876, plain, (spl4_64 <=> (e3 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl4_64])])).
fof(f12169, plain, (~ (e3 = op(e2, e4)) | ~ spl4_64), inference(backward_demodulation, [], [f434, f878])).
fof(f878, plain, ((e3 = op(e2, e2)) | ~ spl4_64), inference(avatar_component_clause, [], [f876])).
fof(f434, plain, ~ (op(e2, e2) = op(e2, e4)), inference(cnf_transformation, [], [f4])).
fof(f12174, plain, (spl4_33 | ~ spl4_64), inference(avatar_split_clause, [], [f12166, f876, f746])).
fof(f12166, plain, ((e2 = op(e3, e3)) | ~ spl4_64), inference(backward_demodulation, [], [f343, f878])).
fof(f12165, plain, (~ spl4_62 | ~ spl4_72), inference(avatar_split_clause, [], [f12157, f910, f868])).
fof(f868, plain, (spl4_62 <=> (e1 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl4_62])])).
fof(f12157, plain, (~ (e1 = op(e2, e2)) | ~ spl4_72), inference(backward_demodulation, [], [f427, f912])).
fof(f427, plain, ~ (op(e2, e0) = op(e2, e2)), inference(cnf_transformation, [], [f4])).
fof(f12154, plain, (spl4_27 | ~ spl4_19 | ~ spl4_80), inference(avatar_split_clause, [], [f12151, f943, f687, f721])).
fof(f12151, plain, ((e1 = op(e3, e4)) | (~ spl4_19 | ~ spl4_80)), inference(backward_demodulation, [], [f12007, f945])).
fof(f12007, plain, ((e1 = op(e3, op(e1, e4))) | ~ spl4_19), inference(forward_demodulation, [], [f340, f689])).
fof(f689, plain, ((e3 = op(e4, e1)) | ~ spl4_19), inference(avatar_component_clause, [], [f687])).
fof(f340, plain, (e1 = op(op(e4, e1), op(e1, e4))), inference(cnf_transformation, [], [f3])).
fof(f12153, plain, (spl4_10 | ~ spl4_19 | ~ spl4_80), inference(avatar_split_clause, [], [f12150, f943, f687, f649])).
fof(f12150, plain, ((e4 = op(e4, e3)) | (~ spl4_19 | ~ spl4_80)), inference(backward_demodulation, [], [f12006, f945])).
fof(f12006, plain, ((e4 = op(op(e1, e4), e3)) | ~ spl4_19), inference(forward_demodulation, [], [f352, f689])).
fof(f352, plain, (e4 = op(op(e1, e4), op(e4, e1))), inference(cnf_transformation, [], [f3])).
fof(f12152, plain, (~ spl4_30 | ~ spl4_80), inference(avatar_split_clause, [], [f12149, f943, f733])).
fof(f12149, plain, (~ (e4 = op(e3, e4)) | ~ spl4_80), inference(backward_demodulation, [], [f401, f945])).
fof(f401, plain, ~ (op(e1, e4) = op(e3, e4)), inference(cnf_transformation, [], [f4])).
fof(f12147, plain, (~ spl4_34 | ~ spl4_84), inference(avatar_split_clause, [], [f12143, f960, f750])).
fof(f12143, plain, (~ (e3 = op(e3, e3)) | ~ spl4_84), inference(backward_demodulation, [], [f391, f962])).
fof(f12140, plain, (spl4_12 | ~ spl4_70 | ~ spl4_88), inference(avatar_split_clause, [], [f12137, f977, f901, f658])).
fof(f977, plain, (spl4_88 <=> (e2 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl4_88])])).
fof(f12137, plain, ((e1 = op(e4, e2)) | (~ spl4_70 | ~ spl4_88)), inference(backward_demodulation, [], [f12009, f979])).
fof(f979, plain, ((e2 = op(e1, e2)) | ~ spl4_88), inference(avatar_component_clause, [], [f977])).
fof(f12009, plain, ((e1 = op(e4, op(e1, e2))) | ~ spl4_70), inference(forward_demodulation, [], [f338, f903])).
fof(f12139, plain, (spl4_53 | ~ spl4_70 | ~ spl4_88), inference(avatar_split_clause, [], [f12136, f977, f901, f830])).
fof(f830, plain, (spl4_53 <=> (e2 = op(e2, e4))), introduced(avatar_definition, [new_symbols(naming, [spl4_53])])).
fof(f12136, plain, ((e2 = op(e2, e4)) | (~ spl4_70 | ~ spl4_88)), inference(backward_demodulation, [], [f12008, f979])).
fof(f12008, plain, ((e2 = op(op(e1, e2), e4)) | ~ spl4_70), inference(forward_demodulation, [], [f342, f903])).
fof(f342, plain, (e2 = op(op(e1, e2), op(e2, e1))), inference(cnf_transformation, [], [f3])).
fof(f12138, plain, (~ spl4_38 | ~ spl4_88), inference(avatar_split_clause, [], [f12130, f977, f767])).
fof(f767, plain, (spl4_38 <=> (e2 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl4_38])])).
fof(f12130, plain, (~ (e2 = op(e3, e2)) | ~ spl4_88), inference(backward_demodulation, [], [f381, f979])).
fof(f381, plain, ~ (op(e1, e2) = op(e3, e2)), inference(cnf_transformation, [], [f4])).
fof(f12126, plain, (~ spl4_77 | ~ spl4_92), inference(avatar_split_clause, [], [f12120, f994, f931])).
fof(f931, plain, (spl4_77 <=> (e1 = op(e1, e4))), introduced(avatar_definition, [new_symbols(naming, [spl4_77])])).
fof(f994, plain, (spl4_92 <=> (e1 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl4_92])])).
fof(f12120, plain, (~ (e1 = op(e1, e4)) | ~ spl4_92), inference(backward_demodulation, [], [f422, f996])).
fof(f996, plain, ((e1 = op(e1, e1)) | ~ spl4_92), inference(avatar_component_clause, [], [f994])).
fof(f422, plain, ~ (op(e1, e1) = op(e1, e4)), inference(cnf_transformation, [], [f4])).
fof(f12125, plain, (~ spl4_42 | ~ spl4_92), inference(avatar_split_clause, [], [f12117, f994, f784])).
fof(f784, plain, (spl4_42 <=> (e1 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl4_42])])).
fof(f12117, plain, (~ (e1 = op(e3, e1)) | ~ spl4_92), inference(backward_demodulation, [], [f371, f996])).
fof(f371, plain, ~ (op(e1, e1) = op(e3, e1)), inference(cnf_transformation, [], [f4])).
fof(f12114, plain, (~ spl4_86 | ~ spl4_96), inference(avatar_split_clause, [], [f12104, f1011, f969])).
fof(f969, plain, (spl4_86 <=> (e0 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl4_86])])).
fof(f12104, plain, (~ (e0 = op(e1, e2)) | ~ spl4_96), inference(backward_demodulation, [], [f417, f1013])).
fof(f417, plain, ~ (op(e1, e0) = op(e1, e2)), inference(cnf_transformation, [], [f4])).
fof(f12113, plain, (~ spl4_46 | ~ spl4_96), inference(avatar_split_clause, [], [f12101, f1011, f801])).
fof(f801, plain, (spl4_46 <=> (e0 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl4_46])])).
fof(f12101, plain, (~ (e0 = op(e3, e0)) | ~ spl4_96), inference(backward_demodulation, [], [f361, f1013])).
fof(f12093, plain, (~ spl4_54 | ~ spl4_104), inference(avatar_split_clause, [], [f12086, f1044, f834])).
fof(f12086, plain, (~ (e3 = op(e2, e4)) | ~ spl4_104), inference(backward_demodulation, [], [f397, f1046])).
fof(f1046, plain, ((e3 = op(e0, e4)) | ~ spl4_104), inference(avatar_component_clause, [], [f1044])).
fof(f397, plain, ~ (op(e0, e4) = op(e2, e4)), inference(cnf_transformation, [], [f4])).
fof(f12082, plain, (~ spl4_86 | ~ spl4_111), inference(avatar_split_clause, [], [f12073, f1074, f969])).
fof(f1074, plain, (spl4_111 <=> (e0 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl4_111])])).
fof(f12073, plain, (~ (e0 = op(e1, e2)) | ~ spl4_111), inference(backward_demodulation, [], [f376, f1076])).
fof(f1076, plain, ((e0 = op(e0, e2)) | ~ spl4_111), inference(avatar_component_clause, [], [f1074])).
fof(f376, plain, ~ (op(e0, e2) = op(e1, e2)), inference(cnf_transformation, [], [f4])).
fof(f12070, plain, (~ spl4_103 | ~ spl4_118), inference(avatar_split_clause, [], [f12066, f1103, f1040])).
fof(f1040, plain, (spl4_103 <=> (e2 = op(e0, e4))), introduced(avatar_definition, [new_symbols(naming, [spl4_103])])).
fof(f1103, plain, (spl4_118 <=> (e2 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl4_118])])).
fof(f12066, plain, (~ (e2 = op(e0, e4)) | ~ spl4_118), inference(backward_demodulation, [], [f412, f1105])).
fof(f1105, plain, ((e2 = op(e0, e1)) | ~ spl4_118), inference(avatar_component_clause, [], [f1103])).
fof(f12069, plain, (~ spl4_93 | ~ spl4_118), inference(avatar_split_clause, [], [f12063, f1103, f998])).
fof(f998, plain, (spl4_93 <=> (e2 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl4_93])])).
fof(f12063, plain, (~ (e2 = op(e1, e1)) | ~ spl4_118), inference(backward_demodulation, [], [f366, f1105])).
fof(f366, plain, ~ (op(e0, e1) = op(e1, e1)), inference(cnf_transformation, [], [f4])).
fof(f12039, plain, (~ spl4_103 | ~ spl4_125 | spl4_171), inference(avatar_split_clause, [], [f12038, f1432, f1132, f1040])).
fof(f1132, plain, (spl4_125 <=> (op(e0, e0) = e4)), introduced(avatar_definition, [new_symbols(naming, [spl4_125])])).
fof(f1432, plain, (spl4_171 <=> (e2 = op(e0, op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl4_171])])).
fof(f12038, plain, (~ (e2 = op(e0, e4)) | (~ spl4_125 | spl4_171)), inference(forward_demodulation, [], [f1434, f1134])).
fof(f1134, plain, ((op(e0, e0) = e4) | ~ spl4_125), inference(avatar_component_clause, [], [f1132])).
fof(f1434, plain, (~ (e2 = op(e0, op(e0, e0))) | spl4_171), inference(avatar_component_clause, [], [f1432])).
fof(f12002, plain, (~ spl4_55 | ~ spl4_70), inference(avatar_split_clause, [], [f12001, f901, f838])).
fof(f838, plain, (spl4_55 <=> (e4 = op(e2, e4))), introduced(avatar_definition, [new_symbols(naming, [spl4_55])])).
fof(f12001, plain, (~ (e4 = op(e2, e4)) | ~ spl4_70), inference(forward_demodulation, [], [f432, f903])).
fof(f432, plain, ~ (op(e2, e1) = op(e2, e4)), inference(cnf_transformation, [], [f4])).
fof(f11996, plain, (~ spl4_24 | ~ spl4_1 | spl4_135), inference(avatar_split_clause, [], [f11974, f1249, f612, f708])).
fof(f708, plain, (spl4_24 <=> (e3 = op(e4, e0))), introduced(avatar_definition, [new_symbols(naming, [spl4_24])])).
fof(f1249, plain, (spl4_135 <=> (e3 = op(e4, op(e4, e4)))), introduced(avatar_definition, [new_symbols(naming, [spl4_135])])).
fof(f11974, plain, (~ (e3 = op(e4, e0)) | (~ spl4_1 | spl4_135)), inference(backward_demodulation, [], [f1251, f614])).
fof(f1251, plain, (~ (e3 = op(e4, op(e4, e4))) | spl4_135), inference(avatar_component_clause, [], [f1249])).
fof(f11989, plain, (~ spl4_1 | ~ spl4_22 | ~ spl4_93 | spl4_175), inference(avatar_contradiction_clause, [], [f11988])).
fof(f11988, plain, ($false | (~ spl4_1 | ~ spl4_22 | ~ spl4_93 | spl4_175)), inference(subsumption_resolution, [], [f11987, f1000])).
fof(f1000, plain, ((e2 = op(e1, e1)) | ~ spl4_93), inference(avatar_component_clause, [], [f998])).
fof(f11987, plain, (~ (e2 = op(e1, e1)) | (~ spl4_1 | ~ spl4_22 | spl4_175)), inference(forward_demodulation, [], [f11979, f702])).
fof(f11979, plain, (~ (e2 = op(op(e4, e0), op(e4, e0))) | (~ spl4_1 | spl4_175)), inference(backward_demodulation, [], [f1459, f614])).
fof(f1459, plain, (~ (e2 = op(op(e4, op(e4, e4)), op(e4, op(e4, e4)))) | spl4_175), inference(avatar_component_clause, [], [f1457])).
fof(f1457, plain, (spl4_175 <=> (e2 = op(op(e4, op(e4, e4)), op(e4, op(e4, e4))))), introduced(avatar_definition, [new_symbols(naming, [spl4_175])])).
fof(f11985, plain, (~ spl4_1 | ~ spl4_22 | spl4_147), inference(avatar_contradiction_clause, [], [f11984])).
fof(f11984, plain, ($false | (~ spl4_1 | ~ spl4_22 | spl4_147)), inference(subsumption_resolution, [], [f11976, f702])).
fof(f11976, plain, (~ (e1 = op(e4, e0)) | (~ spl4_1 | spl4_147)), inference(backward_demodulation, [], [f1307, f614])).
fof(f1307, plain, (~ (e1 = op(e4, op(e4, e4))) | spl4_147), inference(avatar_component_clause, [], [f1305])).
fof(f1305, plain, (spl4_147 <=> (e1 = op(e4, op(e4, e4)))), introduced(avatar_definition, [new_symbols(naming, [spl4_147])])).
fof(f11967, plain, (~ spl4_3 | ~ spl4_8), inference(avatar_split_clause, [], [f11962, f641, f620])).
fof(f11962, plain, (~ (e2 = op(e4, e4)) | ~ spl4_8), inference(backward_demodulation, [], [f455, f643])).
fof(f643, plain, ((e2 = op(e4, e3)) | ~ spl4_8), inference(avatar_component_clause, [], [f641])).
fof(f11961, plain, (spl4_8 | ~ spl4_15 | ~ spl4_54), inference(avatar_split_clause, [], [f11958, f834, f670, f641])).
fof(f11958, plain, ((e2 = op(e4, e3)) | (~ spl4_15 | ~ spl4_54)), inference(backward_demodulation, [], [f11651, f672])).
fof(f672, plain, ((e4 = op(e4, e2)) | ~ spl4_15), inference(avatar_component_clause, [], [f670])).
fof(f11651, plain, ((e2 = op(op(e4, e2), e3)) | ~ spl4_54), inference(backward_demodulation, [], [f345, f836])).
fof(f836, plain, ((e3 = op(e2, e4)) | ~ spl4_54), inference(avatar_component_clause, [], [f834])).
fof(f345, plain, (e2 = op(op(e4, e2), op(e2, e4))), inference(cnf_transformation, [], [f3])).
fof(f11960, plain, (~ spl4_10 | ~ spl4_15), inference(avatar_split_clause, [], [f11956, f670, f649])).
fof(f11956, plain, (~ (e4 = op(e4, e3)) | ~ spl4_15), inference(backward_demodulation, [], [f453, f672])).
fof(f453, plain, ~ (op(e4, e2) = op(e4, e3)), inference(cnf_transformation, [], [f4])).
fof(f11916, plain, (~ spl4_28 | ~ spl4_38), inference(avatar_split_clause, [], [f11911, f767, f725])).
fof(f11911, plain, (~ (e2 = op(e3, e4)) | ~ spl4_38), inference(backward_demodulation, [], [f444, f769])).
fof(f769, plain, ((e2 = op(e3, e2)) | ~ spl4_38), inference(avatar_component_clause, [], [f767])).
fof(f11908, plain, (spl4_19 | ~ spl4_42 | ~ spl4_85), inference(avatar_split_clause, [], [f11906, f964, f784, f687])).
fof(f964, plain, (spl4_85 <=> (e4 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl4_85])])).
fof(f11906, plain, ((e3 = op(e4, e1)) | (~ spl4_42 | ~ spl4_85)), inference(backward_demodulation, [], [f11852, f786])).
fof(f786, plain, ((e1 = op(e3, e1)) | ~ spl4_42), inference(avatar_component_clause, [], [f784])).
fof(f11852, plain, ((e3 = op(e4, op(e3, e1))) | ~ spl4_85), inference(backward_demodulation, [], [f347, f966])).
fof(f966, plain, ((e4 = op(e1, e3)) | ~ spl4_85), inference(avatar_component_clause, [], [f964])).
fof(f11907, plain, (~ spl4_32 | ~ spl4_42), inference(avatar_split_clause, [], [f11900, f784, f742])).
fof(f11900, plain, (~ (e1 = op(e3, e3)) | ~ spl4_42), inference(backward_demodulation, [], [f441, f786])).
fof(f441, plain, ~ (op(e3, e1) = op(e3, e3)), inference(cnf_transformation, [], [f4])).
fof(f11897, plain, (~ spl4_26 | ~ spl4_46), inference(avatar_split_clause, [], [f11890, f801, f717])).
fof(f11890, plain, (~ (e0 = op(e3, e4)) | ~ spl4_46), inference(backward_demodulation, [], [f439, f803])).
fof(f803, plain, ((e0 = op(e3, e0)) | ~ spl4_46), inference(avatar_component_clause, [], [f801])).
fof(f439, plain, ~ (op(e3, e0) = op(e3, e4)), inference(cnf_transformation, [], [f4])).
fof(f11883, plain, (~ spl4_12 | ~ spl4_62), inference(avatar_split_clause, [], [f11877, f868, f658])).
fof(f11877, plain, (~ (e1 = op(e4, e2)) | ~ spl4_62), inference(backward_demodulation, [], [f384, f870])).
fof(f870, plain, ((e1 = op(e2, e2)) | ~ spl4_62), inference(avatar_component_clause, [], [f868])).
fof(f11874, plain, (spl4_22 | ~ spl4_70 | ~ spl4_86), inference(avatar_split_clause, [], [f11870, f969, f901, f700])).
fof(f11870, plain, ((e1 = op(e4, e0)) | (~ spl4_70 | ~ spl4_86)), inference(backward_demodulation, [], [f11735, f903])).
fof(f11735, plain, ((e1 = op(op(e2, e1), e0)) | ~ spl4_86), inference(forward_demodulation, [], [f338, f971])).
fof(f971, plain, ((e0 = op(e1, e2)) | ~ spl4_86), inference(avatar_component_clause, [], [f969])).
fof(f11872, plain, (~ spl4_45 | ~ spl4_70), inference(avatar_split_clause, [], [f11867, f901, f796])).
fof(f796, plain, (spl4_45 <=> (e4 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl4_45])])).
fof(f11867, plain, (~ (e4 = op(e3, e1)) | ~ spl4_70), inference(backward_demodulation, [], [f373, f903])).
fof(f11855, plain, (~ spl4_10 | ~ spl4_85), inference(avatar_split_clause, [], [f11854, f964, f649])).
fof(f11854, plain, (~ (e4 = op(e4, e3)) | ~ spl4_85), inference(backward_demodulation, [], [f392, f966])).
fof(f392, plain, ~ (op(e1, e3) = op(e4, e3)), inference(cnf_transformation, [], [f4])).
fof(f11850, plain, (~ spl4_83 | ~ spl4_93), inference(avatar_split_clause, [], [f11844, f998, f956])).
fof(f956, plain, (spl4_83 <=> (e2 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl4_83])])).
fof(f11844, plain, (~ (e2 = op(e1, e3)) | ~ spl4_93), inference(backward_demodulation, [], [f421, f1000])).
fof(f11849, plain, (~ spl4_68 | ~ spl4_93), inference(avatar_split_clause, [], [f11841, f998, f893])).
fof(f893, plain, (spl4_68 <=> (e2 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl4_68])])).
fof(f11841, plain, (~ (e2 = op(e2, e1)) | ~ spl4_93), inference(backward_demodulation, [], [f370, f1000])).
fof(f370, plain, ~ (op(e1, e1) = op(e2, e1)), inference(cnf_transformation, [], [f4])).
fof(f11848, plain, (spl4_62 | ~ spl4_93), inference(avatar_split_clause, [], [f11840, f998, f868])).
fof(f11840, plain, ((e1 = op(e2, e2)) | ~ spl4_93), inference(backward_demodulation, [], [f337, f1000])).
fof(f11839, plain, (spl4_46 | ~ spl4_99 | ~ spl4_116), inference(avatar_split_clause, [], [f11835, f1095, f1023, f801])).
fof(f1023, plain, (spl4_99 <=> (e3 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl4_99])])).
fof(f11835, plain, ((e0 = op(e3, e0)) | (~ spl4_99 | ~ spl4_116)), inference(backward_demodulation, [], [f11787, f1025])).
fof(f1025, plain, ((e3 = op(e1, e0)) | ~ spl4_99), inference(avatar_component_clause, [], [f1023])).
fof(f11787, plain, ((e0 = op(op(e1, e0), e0)) | ~ spl4_116), inference(backward_demodulation, [], [f332, f1097])).
fof(f11838, plain, (~ spl4_94 | ~ spl4_99), inference(avatar_split_clause, [], [f11832, f1023, f1002])).
fof(f11832, plain, (~ (e3 = op(e1, e1)) | ~ spl4_99), inference(backward_demodulation, [], [f416, f1025])).
fof(f416, plain, ~ (op(e1, e0) = op(e1, e1)), inference(cnf_transformation, [], [f4])).
fof(f11837, plain, (~ spl4_24 | ~ spl4_99), inference(avatar_split_clause, [], [f11831, f1023, f708])).
fof(f11831, plain, (~ (e3 = op(e4, e0)) | ~ spl4_99), inference(backward_demodulation, [], [f362, f1025])).
fof(f362, plain, ~ (op(e1, e0) = op(e4, e0)), inference(cnf_transformation, [], [f4])).
fof(f11828, plain, (~ spl4_3 | ~ spl4_103), inference(avatar_split_clause, [], [f11824, f1040, f620])).
fof(f11824, plain, (~ (e2 = op(e4, e4)) | ~ spl4_103), inference(backward_demodulation, [], [f399, f1042])).
fof(f1042, plain, ((e2 = op(e0, e4)) | ~ spl4_103), inference(avatar_component_clause, [], [f1040])).
fof(f11827, plain, (~ spl4_28 | ~ spl4_103), inference(avatar_split_clause, [], [f11823, f1040, f725])).
fof(f11823, plain, (~ (e2 = op(e3, e4)) | ~ spl4_103), inference(backward_demodulation, [], [f398, f1042])).
fof(f11820, plain, (~ spl4_32 | ~ spl4_107), inference(avatar_split_clause, [], [f11814, f1057, f742])).
fof(f11814, plain, (~ (e1 = op(e3, e3)) | ~ spl4_107), inference(backward_demodulation, [], [f388, f1059])).
fof(f11800, plain, (~ spl4_16 | ~ spl4_116), inference(avatar_split_clause, [], [f11792, f1095, f675])).
fof(f675, plain, (spl4_16 <=> (e0 = op(e4, e1))), introduced(avatar_definition, [new_symbols(naming, [spl4_16])])).
fof(f11792, plain, (~ (e0 = op(e4, e1)) | ~ spl4_116), inference(backward_demodulation, [], [f369, f1097])).
fof(f369, plain, ~ (op(e0, e1) = op(e4, e1)), inference(cnf_transformation, [], [f4])).
fof(f11786, plain, (spl4_103 | ~ spl4_125 | ~ spl4_171), inference(avatar_split_clause, [], [f11781, f1432, f1132, f1040])).
fof(f11781, plain, ((e2 = op(e0, e4)) | (~ spl4_125 | ~ spl4_171)), inference(backward_demodulation, [], [f1433, f1134])).
fof(f1433, plain, ((e2 = op(e0, op(e0, e0))) | ~ spl4_171), inference(avatar_component_clause, [], [f1432])).
fof(f11785, plain, (~ spl4_105 | ~ spl4_125), inference(avatar_split_clause, [], [f11780, f1132, f1048])).
fof(f1048, plain, (spl4_105 <=> (e4 = op(e0, e4))), introduced(avatar_definition, [new_symbols(naming, [spl4_105])])).
fof(f11780, plain, (~ (e4 = op(e0, e4)) | ~ spl4_125), inference(backward_demodulation, [], [f409, f1134])).
fof(f11784, plain, (~ spl4_100 | ~ spl4_125), inference(avatar_split_clause, [], [f11773, f1132, f1027])).
fof(f11773, plain, (~ (e4 = op(e1, e0)) | ~ spl4_125), inference(backward_demodulation, [], [f356, f1134])).
fof(f356, plain, ~ (op(e0, e0) = op(e1, e0)), inference(cnf_transformation, [], [f4])).
fof(f11783, plain, (spl4_1 | ~ spl4_125), inference(avatar_split_clause, [], [f11772, f1132, f612])).
fof(f11772, plain, ((e0 = op(e4, e4)) | ~ spl4_125), inference(backward_demodulation, [], [f331, f1134])).
fof(f11763, plain, (~ spl4_114 | spl4_170 | ~ spl4_171), inference(avatar_split_clause, [], [f11759, f1432, f1426, f1086])).
fof(f1426, plain, (spl4_170 <=> (e3 = op(e0, op(e0, op(e0, e0))))), introduced(avatar_definition, [new_symbols(naming, [spl4_170])])).
fof(f11759, plain, (~ (e3 = op(e0, e2)) | (spl4_170 | ~ spl4_171)), inference(backward_demodulation, [], [f1428, f1433])).
fof(f1428, plain, (~ (e3 = op(e0, op(e0, op(e0, e0)))) | spl4_170), inference(avatar_component_clause, [], [f1426])).
fof(f11720, plain, (~ spl4_3 | ~ spl4_12 | spl4_147), inference(avatar_contradiction_clause, [], [f11719])).
fof(f11719, plain, ($false | (~ spl4_3 | ~ spl4_12 | spl4_147)), inference(subsumption_resolution, [], [f11712, f660])).
fof(f660, plain, ((e1 = op(e4, e2)) | ~ spl4_12), inference(avatar_component_clause, [], [f658])).
fof(f11712, plain, (~ (e1 = op(e4, e2)) | (~ spl4_3 | spl4_147)), inference(backward_demodulation, [], [f1307, f622])).
fof(f622, plain, ((e2 = op(e4, e4)) | ~ spl4_3), inference(avatar_component_clause, [], [f620])).
fof(f11706, plain, (~ spl4_10 | ~ spl4_28 | ~ spl4_54), inference(avatar_contradiction_clause, [], [f11705])).
fof(f11705, plain, ($false | (~ spl4_10 | ~ spl4_28 | ~ spl4_54)), inference(subsumption_resolution, [], [f11704, f465])).
fof(f465, plain, ~ (e3 = e4), inference(cnf_transformation, [], [f5])).
fof(f5, plain, (~ (e3 = e4) & ~ (e2 = e4) & ~ (e2 = e3) & ~ (e1 = e4) & ~ (e1 = e3) & ~ (e1 = e2) & ~ (e0 = e4) & ~ (e0 = e3) & ~ (e0 = e2) & ~ (e0 = e1)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG191+1.p', ax5)).
fof(f11704, plain, ((e3 = e4) | (~ spl4_10 | ~ spl4_28 | ~ spl4_54)), inference(forward_demodulation, [], [f11699, f836])).
fof(f11699, plain, ((e4 = op(e2, e4)) | (~ spl4_10 | ~ spl4_28)), inference(backward_demodulation, [], [f11153, f651])).
fof(f11153, plain, ((e4 = op(e2, op(e4, e3))) | ~ spl4_28), inference(backward_demodulation, [], [f354, f727])).
fof(f727, plain, ((e2 = op(e3, e4)) | ~ spl4_28), inference(avatar_component_clause, [], [f725])).
fof(f354, plain, (e4 = op(op(e3, e4), op(e4, e3))), inference(cnf_transformation, [], [f3])).
fof(f11680, plain, (spl4_10 | ~ spl4_24 | ~ spl4_105), inference(avatar_split_clause, [], [f11678, f1048, f708, f649])).
fof(f11678, plain, ((e4 = op(e4, e3)) | (~ spl4_24 | ~ spl4_105)), inference(backward_demodulation, [], [f11586, f710])).
fof(f710, plain, ((e3 = op(e4, e0)) | ~ spl4_24), inference(avatar_component_clause, [], [f708])).
fof(f11586, plain, ((e4 = op(e4, op(e4, e0))) | ~ spl4_105), inference(backward_demodulation, [], [f351, f1050])).
fof(f1050, plain, ((e4 = op(e0, e4)) | ~ spl4_105), inference(avatar_component_clause, [], [f1048])).
fof(f351, plain, (e4 = op(op(e0, e4), op(e4, e0))), inference(cnf_transformation, [], [f3])).
fof(f11679, plain, (~ spl4_9 | ~ spl4_24), inference(avatar_split_clause, [], [f11674, f708, f645])).
fof(f11674, plain, (~ (e3 = op(e4, e3)) | ~ spl4_24), inference(backward_demodulation, [], [f448, f710])).
fof(f11671, plain, (~ spl4_26 | ~ spl4_28), inference(avatar_contradiction_clause, [], [f11670])).
fof(f11670, plain, ($false | (~ spl4_26 | ~ spl4_28)), inference(subsumption_resolution, [], [f11669, f457])).
fof(f457, plain, ~ (e0 = e2), inference(cnf_transformation, [], [f5])).
fof(f11669, plain, ((e0 = e2) | (~ spl4_26 | ~ spl4_28)), inference(backward_demodulation, [], [f727, f719])).
fof(f719, plain, ((e0 = op(e3, e4)) | ~ spl4_26), inference(avatar_component_clause, [], [f717])).
fof(f11666, plain, (spl4_12 | ~ spl4_45 | ~ spl4_83), inference(avatar_split_clause, [], [f11662, f956, f796, f658])).
fof(f11662, plain, ((e1 = op(e4, e2)) | (~ spl4_45 | ~ spl4_83)), inference(backward_demodulation, [], [f11405, f798])).
fof(f798, plain, ((e4 = op(e3, e1)) | ~ spl4_45), inference(avatar_component_clause, [], [f796])).
fof(f11405, plain, ((e1 = op(op(e3, e1), e2)) | ~ spl4_83), inference(backward_demodulation, [], [f339, f958])).
fof(f958, plain, ((e2 = op(e1, e3)) | ~ spl4_83), inference(avatar_component_clause, [], [f956])).
fof(f11650, plain, (spl4_48 | ~ spl4_39 | ~ spl4_56), inference(avatar_split_clause, [], [f11649, f843, f771, f809])).
fof(f809, plain, (spl4_48 <=> (e2 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl4_48])])).
fof(f11649, plain, ((e2 = op(e3, e0)) | (~ spl4_39 | ~ spl4_56)), inference(backward_demodulation, [], [f11492, f845])).
fof(f11492, plain, ((e2 = op(e3, op(e2, e3))) | ~ spl4_39), inference(forward_demodulation, [], [f344, f773])).
fof(f773, plain, ((e3 = op(e3, e2)) | ~ spl4_39), inference(avatar_component_clause, [], [f771])).
fof(f11641, plain, (spl4_54 | ~ spl4_65 | ~ spl4_140), inference(avatar_split_clause, [], [f11637, f1272, f880, f834])).
fof(f1272, plain, (spl4_140 <=> (e3 = op(e2, op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl4_140])])).
fof(f11637, plain, ((e3 = op(e2, e4)) | (~ spl4_65 | ~ spl4_140)), inference(backward_demodulation, [], [f1273, f882])).
fof(f1273, plain, ((e3 = op(e2, op(e2, e2))) | ~ spl4_140), inference(avatar_component_clause, [], [f1272])).
fof(f11640, plain, (~ spl4_60 | ~ spl4_65), inference(avatar_split_clause, [], [f11635, f880, f859])).
fof(f859, plain, (spl4_60 <=> (e4 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl4_60])])).
fof(f11635, plain, (~ (e4 = op(e2, e3)) | ~ spl4_65), inference(backward_demodulation, [], [f433, f882])).
fof(f433, plain, ~ (op(e2, e2) = op(e2, e3)), inference(cnf_transformation, [], [f4])).
fof(f11632, plain, (~ spl4_52 | ~ spl4_72), inference(avatar_split_clause, [], [f11626, f910, f826])).
fof(f826, plain, (spl4_52 <=> (e1 = op(e2, e4))), introduced(avatar_definition, [new_symbols(naming, [spl4_52])])).
fof(f11626, plain, (~ (e1 = op(e2, e4)) | ~ spl4_72), inference(backward_demodulation, [], [f429, f912])).
fof(f429, plain, ~ (op(e2, e0) = op(e2, e4)), inference(cnf_transformation, [], [f4])).
fof(f11621, plain, (~ spl4_52 | ~ spl4_77), inference(avatar_split_clause, [], [f11616, f931, f826])).
fof(f11616, plain, (~ (e1 = op(e2, e4)) | ~ spl4_77), inference(backward_demodulation, [], [f400, f933])).
fof(f933, plain, ((e1 = op(e1, e4)) | ~ spl4_77), inference(avatar_component_clause, [], [f931])).
fof(f400, plain, ~ (op(e1, e4) = op(e2, e4)), inference(cnf_transformation, [], [f4])).
fof(f11612, plain, (~ spl4_76 | ~ spl4_86), inference(avatar_split_clause, [], [f11604, f969, f927])).
fof(f927, plain, (spl4_76 <=> (e0 = op(e1, e4))), introduced(avatar_definition, [new_symbols(naming, [spl4_76])])).
fof(f11604, plain, (~ (e0 = op(e1, e4)) | ~ spl4_86), inference(backward_demodulation, [], [f424, f971])).
fof(f424, plain, ~ (op(e1, e2) = op(e1, e4)), inference(cnf_transformation, [], [f4])).
fof(f11611, plain, (~ spl4_61 | ~ spl4_86), inference(avatar_split_clause, [], [f11602, f969, f864])).
fof(f11602, plain, (~ (e0 = op(e2, e2)) | ~ spl4_86), inference(backward_demodulation, [], [f380, f971])).
fof(f380, plain, ~ (op(e1, e2) = op(e2, e2)), inference(cnf_transformation, [], [f4])).
fof(f11601, plain, (spl4_77 | ~ spl4_100 | ~ spl4_117), inference(avatar_split_clause, [], [f11597, f1099, f1027, f931])).
fof(f1099, plain, (spl4_117 <=> (e1 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl4_117])])).
fof(f11597, plain, ((e1 = op(e1, e4)) | (~ spl4_100 | ~ spl4_117)), inference(backward_demodulation, [], [f11556, f1029])).
fof(f11556, plain, ((e1 = op(e1, op(e1, e0))) | ~ spl4_117), inference(backward_demodulation, [], [f336, f1101])).
fof(f1101, plain, ((e1 = op(e0, e1)) | ~ spl4_117), inference(avatar_component_clause, [], [f1099])).
fof(f11600, plain, (spl4_16 | ~ spl4_100 | ~ spl4_117), inference(avatar_split_clause, [], [f11596, f1099, f1027, f675])).
fof(f11596, plain, ((e0 = op(e4, e1)) | (~ spl4_100 | ~ spl4_117)), inference(backward_demodulation, [], [f11555, f1029])).
fof(f11555, plain, ((e0 = op(op(e1, e0), e1)) | ~ spl4_117), inference(backward_demodulation, [], [f332, f1101])).
fof(f11599, plain, (~ spl4_90 | ~ spl4_100), inference(avatar_split_clause, [], [f11594, f1027, f985])).
fof(f985, plain, (spl4_90 <=> (e4 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl4_90])])).
fof(f11594, plain, (~ (e4 = op(e1, e2)) | ~ spl4_100), inference(backward_demodulation, [], [f417, f1029])).
fof(f11590, plain, (~ spl4_5 | ~ spl4_105), inference(avatar_split_clause, [], [f11589, f1048, f628])).
fof(f11589, plain, (~ (e4 = op(e4, e4)) | ~ spl4_105), inference(backward_demodulation, [], [f399, f1050])).
fof(f11583, plain, (~ spl4_9 | ~ spl4_109), inference(avatar_split_clause, [], [f11580, f1065, f645])).
fof(f1065, plain, (spl4_109 <=> (e3 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl4_109])])).
fof(f11580, plain, (~ (e3 = op(e4, e3)) | ~ spl4_109), inference(backward_demodulation, [], [f389, f1067])).
fof(f1067, plain, ((e3 = op(e0, e3)) | ~ spl4_109), inference(avatar_component_clause, [], [f1065])).
fof(f11576, plain, (~ spl4_13 | ~ spl4_113), inference(avatar_split_clause, [], [f11571, f1082, f662])).
fof(f1082, plain, (spl4_113 <=> (e2 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl4_113])])).
fof(f11571, plain, (~ (e2 = op(e4, e2)) | ~ spl4_113), inference(backward_demodulation, [], [f379, f1084])).
fof(f1084, plain, ((e2 = op(e0, e2)) | ~ spl4_113), inference(avatar_component_clause, [], [f1082])).
fof(f11565, plain, (~ spl4_17 | ~ spl4_117), inference(avatar_split_clause, [], [f11558, f1099, f679])).
fof(f11558, plain, (~ (e1 = op(e4, e1)) | ~ spl4_117), inference(backward_demodulation, [], [f369, f1101])).
fof(f11543, plain, (~ spl4_106 | ~ spl4_121), inference(avatar_split_clause, [], [f11529, f1116, f1053])).
fof(f1053, plain, (spl4_106 <=> (e0 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl4_106])])).
fof(f1116, plain, (spl4_121 <=> (e0 = op(e0, e0))), introduced(avatar_definition, [new_symbols(naming, [spl4_121])])).
fof(f11529, plain, (~ (e0 = op(e0, e3)) | ~ spl4_121), inference(backward_demodulation, [], [f408, f1118])).
fof(f1118, plain, ((e0 = op(e0, e0)) | ~ spl4_121), inference(avatar_component_clause, [], [f1116])).
fof(f11542, plain, (~ spl4_21 | ~ spl4_121), inference(avatar_split_clause, [], [f11526, f1116, f696])).
fof(f11526, plain, (~ (e0 = op(e4, e0)) | ~ spl4_121), inference(backward_demodulation, [], [f359, f1118])).
fof(f359, plain, ~ (op(e0, e0) = op(e4, e0)), inference(cnf_transformation, [], [f4])).
fof(f11513, plain, (~ spl4_32 | ~ spl4_140 | spl4_162), inference(avatar_contradiction_clause, [], [f11512])).
fof(f11512, plain, ($false | (~ spl4_32 | ~ spl4_140 | spl4_162)), inference(subsumption_resolution, [], [f11511, f744])).
fof(f11511, plain, (~ (e1 = op(e3, e3)) | (~ spl4_140 | spl4_162)), inference(forward_demodulation, [], [f1384, f1273])).
fof(f1384, plain, (~ (e1 = op(op(e2, op(e2, e2)), op(e2, op(e2, e2)))) | spl4_162), inference(avatar_component_clause, [], [f1382])).
fof(f1382, plain, (spl4_162 <=> (e1 = op(op(e2, op(e2, e2)), op(e2, op(e2, e2))))), introduced(avatar_definition, [new_symbols(naming, [spl4_162])])).
fof(f11504, plain, (~ spl4_121 | spl4_130 | ~ spl4_169), inference(avatar_split_clause, [], [f11308, f1419, f1228, f1116])).
fof(f1228, plain, (spl4_130 <=> (e0 = op(op(e3, op(e3, e3)), op(e3, op(e3, e3))))), introduced(avatar_definition, [new_symbols(naming, [spl4_130])])).
fof(f1419, plain, (spl4_169 <=> (e0 = op(e3, op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl4_169])])).
fof(f11308, plain, (~ (e0 = op(e0, e0)) | (spl4_130 | ~ spl4_169)), inference(backward_demodulation, [], [f1230, f1420])).
fof(f1420, plain, ((e0 = op(e3, op(e3, e3))) | ~ spl4_169), inference(avatar_component_clause, [], [f1419])).
fof(f1230, plain, (~ (e0 = op(op(e3, op(e3, e3)), op(e3, op(e3, e3)))) | spl4_130), inference(avatar_component_clause, [], [f1228])).
fof(f11494, plain, (~ spl4_90 | ~ spl4_83 | ~ spl4_94 | spl4_157), inference(avatar_split_clause, [], [f11412, f1356, f1002, f956, f985])).
fof(f1356, plain, (spl4_157 <=> (e4 = op(e1, op(e1, op(e1, e1))))), introduced(avatar_definition, [new_symbols(naming, [spl4_157])])).
fof(f11412, plain, (~ (e4 = op(e1, e2)) | (~ spl4_83 | ~ spl4_94 | spl4_157)), inference(backward_demodulation, [], [f11392, f958])).
fof(f11392, plain, (~ (e4 = op(e1, op(e1, e3))) | (~ spl4_94 | spl4_157)), inference(backward_demodulation, [], [f1358, f1004])).
fof(f1358, plain, (~ (e4 = op(e1, op(e1, op(e1, e1)))) | spl4_157), inference(avatar_component_clause, [], [f1356])).
fof(f11489, plain, (~ spl4_60 | ~ spl4_140 | spl4_155), inference(avatar_split_clause, [], [f11261, f1346, f1272, f859])).
fof(f11261, plain, (~ (e4 = op(e2, e3)) | (~ spl4_140 | spl4_155)), inference(forward_demodulation, [], [f1348, f1273])).
fof(f11485, plain, (~ spl4_48 | ~ spl4_28), inference(avatar_split_clause, [], [f11252, f725, f809])).
fof(f11252, plain, (~ (e2 = op(e3, e0)) | ~ spl4_28), inference(forward_demodulation, [], [f439, f727])).
fof(f11482, plain, (~ spl4_18 | ~ spl4_68), inference(avatar_split_clause, [], [f11481, f893, f683])).
fof(f683, plain, (spl4_18 <=> (e2 = op(e4, e1))), introduced(avatar_definition, [new_symbols(naming, [spl4_18])])).
fof(f11481, plain, (~ (e2 = op(e4, e1)) | ~ spl4_68), inference(forward_demodulation, [], [f374, f895])).
fof(f895, plain, ((e2 = op(e2, e1)) | ~ spl4_68), inference(avatar_component_clause, [], [f893])).
fof(f374, plain, ~ (op(e2, e1) = op(e4, e1)), inference(cnf_transformation, [], [f4])).
fof(f11476, plain, (~ spl4_3 | ~ spl4_28), inference(avatar_split_clause, [], [f11154, f725, f620])).
fof(f11154, plain, (~ (e2 = op(e4, e4)) | ~ spl4_28), inference(backward_demodulation, [], [f405, f727])).
fof(f11470, plain, (~ spl4_31 | ~ spl4_41), inference(avatar_split_clause, [], [f11463, f780, f738])).
fof(f11463, plain, (~ (e0 = op(e3, e3)) | ~ spl4_41), inference(backward_demodulation, [], [f441, f782])).
fof(f782, plain, ((e0 = op(e3, e1)) | ~ spl4_41), inference(avatar_component_clause, [], [f780])).
fof(f11462, plain, (~ spl4_45 | ~ spl4_50), inference(avatar_split_clause, [], [f11458, f817, f796])).
fof(f11458, plain, (~ (e4 = op(e3, e1)) | ~ spl4_50), inference(backward_demodulation, [], [f436, f819])).
fof(f436, plain, ~ (op(e3, e0) = op(e3, e1)), inference(cnf_transformation, [], [f4])).
fof(f11452, plain, (~ spl4_51 | ~ spl4_61), inference(avatar_split_clause, [], [f11445, f864, f822])).
fof(f822, plain, (spl4_51 <=> (e0 = op(e2, e4))), introduced(avatar_definition, [new_symbols(naming, [spl4_51])])).
fof(f11445, plain, (~ (e0 = op(e2, e4)) | ~ spl4_61), inference(backward_demodulation, [], [f434, f866])).
fof(f11443, plain, (spl4_52 | ~ spl4_68 | ~ spl4_90), inference(avatar_split_clause, [], [f11441, f985, f893, f826])).
fof(f11441, plain, ((e1 = op(e2, e4)) | (~ spl4_68 | ~ spl4_90)), inference(backward_demodulation, [], [f11400, f895])).
fof(f11400, plain, ((e1 = op(op(e2, e1), e4)) | ~ spl4_90), inference(backward_demodulation, [], [f338, f987])).
fof(f987, plain, ((e4 = op(e1, e2)) | ~ spl4_90), inference(avatar_component_clause, [], [f985])).
fof(f11434, plain, (~ spl4_69 | ~ spl4_74), inference(avatar_split_clause, [], [f11428, f918, f897])).
fof(f897, plain, (spl4_69 <=> (e3 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl4_69])])).
fof(f918, plain, (spl4_74 <=> (e3 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl4_74])])).
fof(f11428, plain, (~ (e3 = op(e2, e1)) | ~ spl4_74), inference(backward_demodulation, [], [f426, f920])).
fof(f920, plain, ((e3 = op(e2, e0)) | ~ spl4_74), inference(avatar_component_clause, [], [f918])).
fof(f426, plain, ~ (op(e2, e0) = op(e2, e1)), inference(cnf_transformation, [], [f4])).
fof(f11426, plain, (~ spl4_51 | ~ spl4_76), inference(avatar_split_clause, [], [f11419, f927, f822])).
fof(f11419, plain, (~ (e0 = op(e2, e4)) | ~ spl4_76), inference(backward_demodulation, [], [f400, f929])).
fof(f929, plain, ((e0 = op(e1, e4)) | ~ spl4_76), inference(avatar_component_clause, [], [f927])).
fof(f11399, plain, (~ spl4_83 | ~ spl4_94 | spl4_153), inference(avatar_split_clause, [], [f11391, f1336, f1002, f956])).
fof(f1336, plain, (spl4_153 <=> (e2 = op(e1, op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl4_153])])).
fof(f11391, plain, (~ (e2 = op(e1, e3)) | (~ spl4_94 | spl4_153)), inference(backward_demodulation, [], [f1338, f1004])).
fof(f1338, plain, (~ (e2 = op(e1, op(e1, e1))) | spl4_153), inference(avatar_component_clause, [], [f1336])).
fof(f11398, plain, (~ spl4_79 | ~ spl4_94), inference(avatar_split_clause, [], [f11386, f1002, f939])).
fof(f939, plain, (spl4_79 <=> (e3 = op(e1, e4))), introduced(avatar_definition, [new_symbols(naming, [spl4_79])])).
fof(f11386, plain, (~ (e3 = op(e1, e4)) | ~ spl4_94), inference(backward_demodulation, [], [f422, f1004])).
fof(f11397, plain, (~ spl4_69 | ~ spl4_94), inference(avatar_split_clause, [], [f11382, f1002, f897])).
fof(f11382, plain, (~ (e3 = op(e2, e1)) | ~ spl4_94), inference(backward_demodulation, [], [f370, f1004])).
fof(f11380, plain, (spl4_76 | ~ spl4_97 | ~ spl4_120), inference(avatar_split_clause, [], [f11376, f1111, f1015, f927])).
fof(f1015, plain, (spl4_97 <=> (e1 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl4_97])])).
fof(f11376, plain, ((e0 = op(e1, e4)) | (~ spl4_97 | ~ spl4_120)), inference(backward_demodulation, [], [f11329, f1017])).
fof(f1017, plain, ((e1 = op(e1, e0)) | ~ spl4_97), inference(avatar_component_clause, [], [f1015])).
fof(f11329, plain, ((e0 = op(op(e1, e0), e4)) | ~ spl4_120), inference(backward_demodulation, [], [f332, f1113])).
fof(f11379, plain, (~ spl4_82 | ~ spl4_97), inference(avatar_split_clause, [], [f11371, f1015, f952])).
fof(f952, plain, (spl4_82 <=> (e1 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl4_82])])).
fof(f11371, plain, (~ (e1 = op(e1, e3)) | ~ spl4_97), inference(backward_demodulation, [], [f418, f1017])).
fof(f418, plain, ~ (op(e1, e0) = op(e1, e3)), inference(cnf_transformation, [], [f4])).
fof(f11378, plain, (~ spl4_47 | ~ spl4_97), inference(avatar_split_clause, [], [f11368, f1015, f805])).
fof(f805, plain, (spl4_47 <=> (e1 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl4_47])])).
fof(f11368, plain, (~ (e1 = op(e3, e0)) | ~ spl4_97), inference(backward_demodulation, [], [f361, f1017])).
fof(f11365, plain, (~ spl4_79 | ~ spl4_104), inference(avatar_split_clause, [], [f11360, f1044, f939])).
fof(f11360, plain, (~ (e3 = op(e1, e4)) | ~ spl4_104), inference(backward_demodulation, [], [f396, f1046])).
fof(f396, plain, ~ (op(e0, e4) = op(e1, e4)), inference(cnf_transformation, [], [f4])).
fof(f11359, plain, (~ spl4_31 | ~ spl4_106), inference(avatar_split_clause, [], [f11353, f1053, f738])).
fof(f11353, plain, (~ (e0 = op(e3, e3)) | ~ spl4_106), inference(backward_demodulation, [], [f388, f1055])).
fof(f1055, plain, ((e0 = op(e0, e3)) | ~ spl4_106), inference(avatar_component_clause, [], [f1053])).
fof(f11349, plain, (~ spl4_102 | ~ spl4_112), inference(avatar_split_clause, [], [f11344, f1078, f1036])).
fof(f11344, plain, (~ (e1 = op(e0, e4)) | ~ spl4_112), inference(backward_demodulation, [], [f414, f1080])).
fof(f414, plain, ~ (op(e0, e2) = op(e0, e4)), inference(cnf_transformation, [], [f4])).
fof(f11348, plain, (~ spl4_62 | ~ spl4_112), inference(avatar_split_clause, [], [f11342, f1078, f868])).
fof(f11342, plain, (~ (e1 = op(e2, e2)) | ~ spl4_112), inference(backward_demodulation, [], [f377, f1080])).
fof(f377, plain, ~ (op(e0, e2) = op(e2, e2)), inference(cnf_transformation, [], [f4])).
fof(f11338, plain, (~ spl4_115 | ~ spl4_120), inference(avatar_split_clause, [], [f11334, f1111, f1090])).
fof(f1090, plain, (spl4_115 <=> (e4 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl4_115])])).
fof(f11334, plain, (~ (e4 = op(e0, e2)) | ~ spl4_120), inference(backward_demodulation, [], [f410, f1113])).
fof(f410, plain, ~ (op(e0, e1) = op(e0, e2)), inference(cnf_transformation, [], [f4])).
fof(f11337, plain, (~ spl4_45 | ~ spl4_120), inference(avatar_split_clause, [], [f11333, f1111, f796])).
fof(f11333, plain, (~ (e4 = op(e3, e1)) | ~ spl4_120), inference(backward_demodulation, [], [f368, f1113])).
fof(f368, plain, ~ (op(e0, e1) = op(e3, e1)), inference(cnf_transformation, [], [f4])).
fof(f11301, plain, (~ spl4_93 | spl4_178 | ~ spl4_180), inference(avatar_split_clause, [], [f11300, f1492, f1476, f998])).
fof(f1476, plain, (spl4_178 <=> (e2 = op(op(e0, op(e0, e0)), op(e0, op(e0, e0))))), introduced(avatar_definition, [new_symbols(naming, [spl4_178])])).
fof(f1492, plain, (spl4_180 <=> (e1 = op(e0, op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl4_180])])).
fof(f11300, plain, (~ (e2 = op(e1, e1)) | (spl4_178 | ~ spl4_180)), inference(forward_demodulation, [], [f1478, f1493])).
fof(f1493, plain, ((e1 = op(e0, op(e0, e0))) | ~ spl4_180), inference(avatar_component_clause, [], [f1492])).
fof(f1478, plain, (~ (e2 = op(op(e0, op(e0, e0)), op(e0, op(e0, e0)))) | spl4_178), inference(avatar_component_clause, [], [f1476])).
fof(f11299, plain, (~ spl4_94 | ~ spl4_180 | spl4_185), inference(avatar_split_clause, [], [f11296, f1528, f1492, f1002])).
fof(f1528, plain, (spl4_185 <=> (e3 = op(op(e0, op(e0, e0)), op(e0, op(e0, e0))))), introduced(avatar_definition, [new_symbols(naming, [spl4_185])])).
fof(f11296, plain, (~ (e3 = op(e1, e1)) | (~ spl4_180 | spl4_185)), inference(backward_demodulation, [], [f1530, f1493])).
fof(f1530, plain, (~ (e3 = op(op(e0, op(e0, e0)), op(e0, op(e0, e0)))) | spl4_185), inference(avatar_component_clause, [], [f1528])).
fof(f11235, plain, (~ spl4_86 | ~ spl4_93 | ~ spl4_100 | spl4_157), inference(avatar_contradiction_clause, [], [f11234])).
fof(f11234, plain, ($false | (~ spl4_86 | ~ spl4_93 | ~ spl4_100 | spl4_157)), inference(subsumption_resolution, [], [f11233, f1029])).
fof(f11233, plain, (~ (e4 = op(e1, e0)) | (~ spl4_86 | ~ spl4_93 | spl4_157)), inference(forward_demodulation, [], [f11232, f971])).
fof(f11232, plain, (~ (e4 = op(e1, op(e1, e2))) | (~ spl4_93 | spl4_157)), inference(forward_demodulation, [], [f1358, f1000])).
fof(f11202, plain, (~ spl4_4 | ~ spl4_5 | spl4_183), inference(avatar_split_clause, [], [f11201, f1513, f628, f624])).
fof(f624, plain, (spl4_4 <=> (e3 = op(e4, e4))), introduced(avatar_definition, [new_symbols(naming, [spl4_4])])).
fof(f1513, plain, (spl4_183 <=> (e3 = op(op(e4, op(e4, e4)), op(e4, op(e4, e4))))), introduced(avatar_definition, [new_symbols(naming, [spl4_183])])).
fof(f11201, plain, (~ (e3 = op(e4, e4)) | (~ spl4_5 | spl4_183)), inference(forward_demodulation, [], [f11196, f630])).
fof(f630, plain, ((e4 = op(e4, e4)) | ~ spl4_5), inference(avatar_component_clause, [], [f628])).
fof(f11196, plain, (~ (e3 = op(op(e4, e4), op(e4, e4))) | (~ spl4_5 | spl4_183)), inference(backward_demodulation, [], [f1515, f630])).
fof(f1515, plain, (~ (e3 = op(op(e4, op(e4, e4)), op(e4, op(e4, e4)))) | spl4_183), inference(avatar_component_clause, [], [f1513])).
fof(f11198, plain, (~ spl4_4 | ~ spl4_5 | spl4_135), inference(avatar_split_clause, [], [f11190, f1249, f628, f624])).
fof(f11190, plain, (~ (e3 = op(e4, e4)) | (~ spl4_5 | spl4_135)), inference(backward_demodulation, [], [f1251, f630])).
fof(f11187, plain, (~ spl4_4 | ~ spl4_9), inference(avatar_split_clause, [], [f11183, f645, f624])).
fof(f11183, plain, (~ (e3 = op(e4, e4)) | ~ spl4_9), inference(backward_demodulation, [], [f455, f647])).
fof(f647, plain, ((e3 = op(e4, e3)) | ~ spl4_9), inference(avatar_component_clause, [], [f645])).
fof(f11176, plain, (~ spl4_7 | ~ spl4_17), inference(avatar_split_clause, [], [f11169, f679, f637])).
fof(f637, plain, (spl4_7 <=> (e1 = op(e4, e3))), introduced(avatar_definition, [new_symbols(naming, [spl4_7])])).
fof(f11169, plain, (~ (e1 = op(e4, e3)) | ~ spl4_17), inference(backward_demodulation, [], [f451, f681])).
fof(f681, plain, ((e1 = op(e4, e1)) | ~ spl4_17), inference(avatar_component_clause, [], [f679])).
fof(f451, plain, ~ (op(e4, e1) = op(e4, e3)), inference(cnf_transformation, [], [f4])).
fof(f11151, plain, (~ spl4_31 | spl4_46 | ~ spl4_169), inference(avatar_contradiction_clause, [], [f11150])).
fof(f11150, plain, ($false | (~ spl4_31 | spl4_46 | ~ spl4_169)), inference(subsumption_resolution, [], [f11148, f802])).
fof(f802, plain, (~ (e0 = op(e3, e0)) | spl4_46), inference(avatar_component_clause, [], [f801])).
fof(f11148, plain, ((e0 = op(e3, e0)) | (~ spl4_31 | ~ spl4_169)), inference(backward_demodulation, [], [f1420, f740])).
fof(f740, plain, ((e0 = op(e3, e3)) | ~ spl4_31), inference(avatar_component_clause, [], [f738])).
fof(f11149, plain, (~ spl4_26 | ~ spl4_31), inference(avatar_split_clause, [], [f11143, f738, f717])).
fof(f11143, plain, (~ (e0 = op(e3, e4)) | ~ spl4_31), inference(backward_demodulation, [], [f445, f740])).
fof(f11140, plain, (spl4_9 | ~ spl4_39 | ~ spl4_60), inference(avatar_split_clause, [], [f11138, f859, f771, f645])).
fof(f11138, plain, ((e3 = op(e4, e3)) | (~ spl4_39 | ~ spl4_60)), inference(backward_demodulation, [], [f11099, f773])).
fof(f11099, plain, ((e3 = op(e4, op(e3, e2))) | ~ spl4_60), inference(backward_demodulation, [], [f348, f861])).
fof(f861, plain, ((e4 = op(e2, e3)) | ~ spl4_60), inference(avatar_component_clause, [], [f859])).
fof(f348, plain, (e3 = op(op(e2, e3), op(e3, e2))), inference(cnf_transformation, [], [f3])).
fof(f11139, plain, (spl4_28 | ~ spl4_39 | ~ spl4_60), inference(avatar_split_clause, [], [f11137, f859, f771, f725])).
fof(f11137, plain, ((e2 = op(e3, e4)) | (~ spl4_39 | ~ spl4_60)), inference(backward_demodulation, [], [f11098, f773])).
fof(f11098, plain, ((e2 = op(op(e3, e2), e4)) | ~ spl4_60), inference(backward_demodulation, [], [f344, f861])).
fof(f11132, plain, (spl4_17 | ~ spl4_45 | ~ spl4_82), inference(avatar_split_clause, [], [f11129, f952, f796, f679])).
fof(f11129, plain, ((e1 = op(e4, e1)) | (~ spl4_45 | ~ spl4_82)), inference(backward_demodulation, [], [f11030, f798])).
fof(f11030, plain, ((e1 = op(op(e3, e1), e1)) | ~ spl4_82), inference(backward_demodulation, [], [f339, f954])).
fof(f954, plain, ((e1 = op(e1, e3)) | ~ spl4_82), inference(avatar_component_clause, [], [f952])).
fof(f11131, plain, (~ spl4_35 | ~ spl4_45), inference(avatar_split_clause, [], [f11127, f796, f754])).
fof(f754, plain, (spl4_35 <=> (e4 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl4_35])])).
fof(f11127, plain, (~ (e4 = op(e3, e3)) | ~ spl4_45), inference(backward_demodulation, [], [f441, f798])).
fof(f11124, plain, (~ spl4_37 | ~ spl4_47), inference(avatar_split_clause, [], [f11116, f805, f763])).
fof(f763, plain, (spl4_37 <=> (e1 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl4_37])])).
fof(f11116, plain, (~ (e1 = op(e3, e2)) | ~ spl4_47), inference(backward_demodulation, [], [f437, f807])).
fof(f807, plain, ((e1 = op(e3, e0)) | ~ spl4_47), inference(avatar_component_clause, [], [f805])).
fof(f11113, plain, (~ spl4_26 | ~ spl4_51), inference(avatar_split_clause, [], [f11107, f822, f717])).
fof(f11107, plain, (~ (e0 = op(e3, e4)) | ~ spl4_51), inference(backward_demodulation, [], [f403, f824])).
fof(f824, plain, ((e0 = op(e2, e4)) | ~ spl4_51), inference(avatar_component_clause, [], [f822])).
fof(f403, plain, ~ (op(e2, e4) = op(e3, e4)), inference(cnf_transformation, [], [f4])).
fof(f11104, plain, (~ spl4_55 | ~ spl4_60), inference(avatar_split_clause, [], [f11102, f859, f838])).
fof(f11102, plain, (~ (e4 = op(e2, e4)) | ~ spl4_60), inference(backward_demodulation, [], [f435, f861])).
fof(f435, plain, ~ (op(e2, e3) = op(e2, e4)), inference(cnf_transformation, [], [f4])).
fof(f11103, plain, (~ spl4_35 | ~ spl4_60), inference(avatar_split_clause, [], [f11100, f859, f754])).
fof(f11100, plain, (~ (e4 = op(e3, e3)) | ~ spl4_60), inference(backward_demodulation, [], [f393, f861])).
fof(f393, plain, ~ (op(e2, e3) = op(e3, e3)), inference(cnf_transformation, [], [f4])).
fof(f11094, plain, (~ spl4_62 | ~ spl4_69 | spl4_140), inference(avatar_contradiction_clause, [], [f11093])).
fof(f11093, plain, ($false | (~ spl4_62 | ~ spl4_69 | spl4_140)), inference(subsumption_resolution, [], [f11083, f899])).
fof(f899, plain, ((e3 = op(e2, e1)) | ~ spl4_69), inference(avatar_component_clause, [], [f897])).
fof(f11083, plain, (~ (e3 = op(e2, e1)) | (~ spl4_62 | spl4_140)), inference(backward_demodulation, [], [f1274, f870])).
fof(f1274, plain, (~ (e3 = op(e2, op(e2, e2))) | spl4_140), inference(avatar_component_clause, [], [f1272])).
fof(f11087, plain, (~ spl4_37 | ~ spl4_62), inference(avatar_split_clause, [], [f11074, f868, f763])).
fof(f11074, plain, (~ (e1 = op(e3, e2)) | ~ spl4_62), inference(backward_demodulation, [], [f383, f870])).
fof(f11072, plain, (spl4_47 | ~ spl4_69 | ~ spl4_86), inference(avatar_split_clause, [], [f11068, f969, f897, f805])).
fof(f11068, plain, ((e1 = op(e3, e0)) | (~ spl4_69 | ~ spl4_86)), inference(backward_demodulation, [], [f11013, f899])).
fof(f11013, plain, ((e1 = op(op(e2, e1), e0)) | ~ spl4_86), inference(backward_demodulation, [], [f338, f971])).
fof(f11070, plain, (~ spl4_44 | ~ spl4_69), inference(avatar_split_clause, [], [f11062, f897, f792])).
fof(f792, plain, (spl4_44 <=> (e3 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl4_44])])).
fof(f11062, plain, (~ (e3 = op(e3, e1)) | ~ spl4_69), inference(backward_demodulation, [], [f373, f899])).
fof(f11061, plain, (spl4_13 | ~ spl4_73 | ~ spl4_115), inference(avatar_split_clause, [], [f11057, f1090, f914, f662])).
fof(f11057, plain, ((e2 = op(e4, e2)) | (~ spl4_73 | ~ spl4_115)), inference(backward_demodulation, [], [f10947, f916])).
fof(f10947, plain, ((e2 = op(e4, op(e2, e0))) | ~ spl4_115), inference(backward_demodulation, [], [f341, f1092])).
fof(f1092, plain, ((e4 = op(e0, e2)) | ~ spl4_115), inference(avatar_component_clause, [], [f1090])).
fof(f341, plain, (e2 = op(op(e0, e2), op(e2, e0))), inference(cnf_transformation, [], [f3])).
fof(f11047, plain, (~ spl4_4 | ~ spl4_79), inference(avatar_split_clause, [], [f11045, f939, f624])).
fof(f11045, plain, (~ (e3 = op(e4, e4)) | ~ spl4_79), inference(backward_demodulation, [], [f402, f941])).
fof(f941, plain, ((e3 = op(e1, e4)) | ~ spl4_79), inference(avatar_component_clause, [], [f939])).
fof(f11040, plain, (~ spl4_77 | ~ spl4_82), inference(avatar_split_clause, [], [f11035, f952, f931])).
fof(f11035, plain, (~ (e1 = op(e1, e4)) | ~ spl4_82), inference(backward_demodulation, [], [f425, f954])).
fof(f425, plain, ~ (op(e1, e3) = op(e1, e4)), inference(cnf_transformation, [], [f4])).
fof(f11039, plain, (~ spl4_7 | ~ spl4_82), inference(avatar_split_clause, [], [f11034, f952, f637])).
fof(f11034, plain, (~ (e1 = op(e4, e3)) | ~ spl4_82), inference(backward_demodulation, [], [f392, f954])).
fof(f11029, plain, (~ spl4_86 | ~ spl4_93 | ~ spl4_124 | spl4_184), inference(avatar_contradiction_clause, [], [f11028])).
fof(f11028, plain, ($false | (~ spl4_86 | ~ spl4_93 | ~ spl4_124 | spl4_184)), inference(subsumption_resolution, [], [f11026, f1130])).
fof(f11026, plain, (~ (op(e0, e0) = e3) | (~ spl4_86 | ~ spl4_93 | spl4_184)), inference(backward_demodulation, [], [f11007, f971])).
fof(f11007, plain, (~ (e3 = op(op(e1, e2), op(e1, e2))) | (~ spl4_93 | spl4_184)), inference(backward_demodulation, [], [f1520, f1000])).
fof(f1520, plain, (~ (e3 = op(op(e1, op(e1, e1)), op(e1, op(e1, e1)))) | spl4_184), inference(avatar_component_clause, [], [f1518])).
fof(f1518, plain, (spl4_184 <=> (e3 = op(op(e1, op(e1, e1)), op(e1, op(e1, e1))))), introduced(avatar_definition, [new_symbols(naming, [spl4_184])])).
fof(f11027, plain, (~ spl4_11 | ~ spl4_86), inference(avatar_split_clause, [], [f11017, f969, f654])).
fof(f11017, plain, (~ (e0 = op(e4, e2)) | ~ spl4_86), inference(backward_demodulation, [], [f382, f971])).
fof(f382, plain, ~ (op(e1, e2) = op(e4, e2)), inference(cnf_transformation, [], [f4])).
fof(f11012, plain, (~ spl4_86 | ~ spl4_93 | spl4_181), inference(avatar_split_clause, [], [f11006, f1497, f998, f969])).
fof(f1497, plain, (spl4_181 <=> (e0 = op(e1, op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl4_181])])).
fof(f11006, plain, (~ (e0 = op(e1, e2)) | (~ spl4_93 | spl4_181)), inference(backward_demodulation, [], [f1499, f1000])).
fof(f1499, plain, (~ (e0 = op(e1, op(e1, e1))) | spl4_181), inference(avatar_component_clause, [], [f1497])).
fof(f11011, plain, (~ spl4_90 | ~ spl4_93 | spl4_146), inference(avatar_split_clause, [], [f11002, f1300, f998, f985])).
fof(f1300, plain, (spl4_146 <=> (e4 = op(e1, op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl4_146])])).
fof(f11002, plain, (~ (e4 = op(e1, e2)) | (~ spl4_93 | spl4_146)), inference(backward_demodulation, [], [f1302, f1000])).
fof(f1302, plain, (~ (e4 = op(e1, op(e1, e1))) | spl4_146), inference(avatar_component_clause, [], [f1300])).
fof(f11009, plain, (~ spl4_18 | ~ spl4_93), inference(avatar_split_clause, [], [f10996, f998, f683])).
fof(f10996, plain, (~ (e2 = op(e4, e1)) | ~ spl4_93), inference(backward_demodulation, [], [f372, f1000])).
fof(f10990, plain, (~ spl4_25 | ~ spl4_100), inference(avatar_split_clause, [], [f10983, f1027, f712])).
fof(f712, plain, (spl4_25 <=> (e4 = op(e4, e0))), introduced(avatar_definition, [new_symbols(naming, [spl4_25])])).
fof(f10983, plain, (~ (e4 = op(e4, e0)) | ~ spl4_100), inference(backward_demodulation, [], [f362, f1029])).
fof(f10980, plain, (~ spl4_77 | ~ spl4_102), inference(avatar_split_clause, [], [f10973, f1036, f931])).
fof(f10973, plain, (~ (e1 = op(e1, e4)) | ~ spl4_102), inference(backward_demodulation, [], [f396, f1038])).
fof(f10970, plain, (~ spl4_103 | ~ spl4_108), inference(avatar_split_clause, [], [f10961, f1061, f1040])).
fof(f10961, plain, (~ (e2 = op(e0, e4)) | ~ spl4_108), inference(backward_demodulation, [], [f415, f1063])).
fof(f415, plain, ~ (op(e0, e3) = op(e0, e4)), inference(cnf_transformation, [], [f4])).
fof(f10969, plain, (~ spl4_83 | ~ spl4_108), inference(avatar_split_clause, [], [f10957, f1061, f956])).
fof(f10957, plain, (~ (e2 = op(e1, e3)) | ~ spl4_108), inference(backward_demodulation, [], [f386, f1063])).
fof(f386, plain, ~ (op(e0, e3) = op(e1, e3)), inference(cnf_transformation, [], [f4])).
fof(f10954, plain, (~ spl4_90 | ~ spl4_115), inference(avatar_split_clause, [], [f10948, f1090, f985])).
fof(f10948, plain, (~ (e4 = op(e1, e2)) | ~ spl4_115), inference(backward_demodulation, [], [f376, f1092])).
fof(f10944, plain, (~ spl4_91 | ~ spl4_116), inference(avatar_split_clause, [], [f10933, f1095, f990])).
fof(f990, plain, (spl4_91 <=> (e0 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl4_91])])).
fof(f10933, plain, (~ (e0 = op(e1, e1)) | ~ spl4_116), inference(backward_demodulation, [], [f366, f1097])).
fof(f10929, plain, (~ spl4_114 | ~ spl4_124), inference(avatar_split_clause, [], [f10915, f1128, f1086])).
fof(f10915, plain, (~ (e3 = op(e0, e2)) | ~ spl4_124), inference(backward_demodulation, [], [f407, f1130])).
fof(f407, plain, ~ (op(e0, e0) = op(e0, e2)), inference(cnf_transformation, [], [f4])).
fof(f10928, plain, (~ spl4_99 | ~ spl4_124), inference(avatar_split_clause, [], [f10910, f1128, f1023])).
fof(f10910, plain, (~ (e3 = op(e1, e0)) | ~ spl4_124), inference(backward_demodulation, [], [f356, f1130])).
fof(f10883, plain, (~ spl4_13 | ~ spl4_18), inference(avatar_split_clause, [], [f10876, f683, f662])).
fof(f10876, plain, (~ (e2 = op(e4, e2)) | ~ spl4_18), inference(backward_demodulation, [], [f450, f685])).
fof(f685, plain, ((e2 = op(e4, e1)) | ~ spl4_18), inference(avatar_component_clause, [], [f683])).
fof(f450, plain, ~ (op(e4, e1) = op(e4, e2)), inference(cnf_transformation, [], [f4])).
fof(f10874, plain, (~ spl4_5 | ~ spl4_25), inference(avatar_split_clause, [], [f10871, f712, f628])).
fof(f10871, plain, (~ (e4 = op(e4, e4)) | ~ spl4_25), inference(backward_demodulation, [], [f449, f714])).
fof(f714, plain, ((e4 = op(e4, e0)) | ~ spl4_25), inference(avatar_component_clause, [], [f712])).
fof(f449, plain, ~ (op(e4, e0) = op(e4, e4)), inference(cnf_transformation, [], [f4])).
fof(f10859, plain, (spl4_4 | ~ spl4_35), inference(avatar_split_clause, [], [f10855, f754, f624])).
fof(f10855, plain, ((e3 = op(e4, e4)) | ~ spl4_35), inference(backward_demodulation, [], [f349, f756])).
fof(f756, plain, ((e4 = op(e3, e3)) | ~ spl4_35), inference(avatar_component_clause, [], [f754])).
fof(f349, plain, (e3 = op(op(e3, e3), op(e3, e3))), inference(cnf_transformation, [], [f3])).
fof(f10845, plain, (~ spl4_5 | ~ spl4_55), inference(avatar_split_clause, [], [f10844, f838, f628])).
fof(f10844, plain, (~ (e4 = op(e4, e4)) | ~ spl4_55), inference(backward_demodulation, [], [f404, f840])).
fof(f840, plain, ((e4 = op(e2, e4)) | ~ spl4_55), inference(avatar_component_clause, [], [f838])).
fof(f404, plain, ~ (op(e2, e4) = op(e4, e4)), inference(cnf_transformation, [], [f4])).
fof(f10840, plain, (~ spl4_54 | ~ spl4_59), inference(avatar_split_clause, [], [f10837, f855, f834])).
fof(f10837, plain, (~ (e3 = op(e2, e4)) | ~ spl4_59), inference(backward_demodulation, [], [f435, f857])).
fof(f857, plain, ((e3 = op(e2, e3)) | ~ spl4_59), inference(avatar_component_clause, [], [f855])).
fof(f10839, plain, (~ spl4_9 | ~ spl4_59), inference(avatar_split_clause, [], [f10836, f855, f645])).
fof(f10836, plain, (~ (e3 = op(e4, e3)) | ~ spl4_59), inference(backward_demodulation, [], [f394, f857])).
fof(f394, plain, ~ (op(e2, e3) = op(e4, e3)), inference(cnf_transformation, [], [f4])).
fof(f10824, plain, (~ spl4_58 | ~ spl4_63), inference(avatar_split_clause, [], [f10814, f872, f851])).
fof(f872, plain, (spl4_63 <=> (e2 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl4_63])])).
fof(f10814, plain, (~ (e2 = op(e2, e3)) | ~ spl4_63), inference(backward_demodulation, [], [f433, f874])).
fof(f874, plain, ((e2 = op(e2, e2)) | ~ spl4_63), inference(avatar_component_clause, [], [f872])).
fof(f10823, plain, (~ spl4_13 | ~ spl4_63), inference(avatar_split_clause, [], [f10813, f872, f662])).
fof(f10813, plain, (~ (e2 = op(e4, e2)) | ~ spl4_63), inference(backward_demodulation, [], [f384, f874])).
fof(f10810, plain, (spl4_18 | ~ spl4_67 | ~ spl4_90), inference(avatar_split_clause, [], [f10807, f985, f889, f683])).
fof(f889, plain, (spl4_67 <=> (e1 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl4_67])])).
fof(f10807, plain, ((e2 = op(e4, e1)) | (~ spl4_67 | ~ spl4_90)), inference(backward_demodulation, [], [f10758, f891])).
fof(f891, plain, ((e1 = op(e2, e1)) | ~ spl4_67), inference(avatar_component_clause, [], [f889])).
fof(f10758, plain, ((e2 = op(e4, op(e2, e1))) | ~ spl4_90), inference(backward_demodulation, [], [f342, f987])).
fof(f10809, plain, (~ spl4_62 | ~ spl4_67), inference(avatar_split_clause, [], [f10800, f889, f868])).
fof(f10800, plain, (~ (e1 = op(e2, e2)) | ~ spl4_67), inference(backward_demodulation, [], [f430, f891])).
fof(f10808, plain, (~ spl4_17 | ~ spl4_67), inference(avatar_split_clause, [], [f10799, f889, f679])).
fof(f10799, plain, (~ (e1 = op(e4, e1)) | ~ spl4_67), inference(backward_demodulation, [], [f374, f891])).
fof(f10785, plain, (~ spl4_27 | ~ spl4_77), inference(avatar_split_clause, [], [f10780, f931, f721])).
fof(f10780, plain, (~ (e1 = op(e3, e4)) | ~ spl4_77), inference(backward_demodulation, [], [f401, f933])).
fof(f10775, plain, (spl4_59 | ~ spl4_44 | ~ spl4_83), inference(avatar_split_clause, [], [f10772, f956, f792, f855])).
fof(f10772, plain, ((e3 = op(e2, e3)) | (~ spl4_44 | ~ spl4_83)), inference(backward_demodulation, [], [f10633, f958])).
fof(f10633, plain, ((e3 = op(op(e1, e3), e3)) | ~ spl4_44), inference(forward_demodulation, [], [f347, f794])).
fof(f794, plain, ((e3 = op(e3, e1)) | ~ spl4_44), inference(avatar_component_clause, [], [f792])).
fof(f10774, plain, (~ spl4_58 | ~ spl4_83), inference(avatar_split_clause, [], [f10766, f956, f851])).
fof(f10766, plain, (~ (e2 = op(e2, e3)) | ~ spl4_83), inference(backward_demodulation, [], [f390, f958])).
fof(f10765, plain, (~ spl4_85 | ~ spl4_90), inference(avatar_split_clause, [], [f10762, f985, f964])).
fof(f10762, plain, (~ (e4 = op(e1, e3)) | ~ spl4_90), inference(backward_demodulation, [], [f423, f987])).
fof(f423, plain, ~ (op(e1, e2) = op(e1, e3)), inference(cnf_transformation, [], [f4])).
fof(f10764, plain, (~ spl4_40 | ~ spl4_90), inference(avatar_split_clause, [], [f10760, f985, f775])).
fof(f10760, plain, (~ (e4 = op(e3, e2)) | ~ spl4_90), inference(backward_demodulation, [], [f381, f987])).
fof(f10756, plain, (~ spl4_76 | ~ spl4_91), inference(avatar_split_clause, [], [f10748, f990, f927])).
fof(f10748, plain, (~ (e0 = op(e1, e4)) | ~ spl4_91), inference(backward_demodulation, [], [f422, f992])).
fof(f992, plain, ((e0 = op(e1, e1)) | ~ spl4_91), inference(avatar_component_clause, [], [f990])).
fof(f10755, plain, (~ spl4_66 | ~ spl4_91), inference(avatar_split_clause, [], [f10744, f990, f885])).
fof(f10744, plain, (~ (e0 = op(e2, e1)) | ~ spl4_91), inference(backward_demodulation, [], [f370, f992])).
fof(f10742, plain, (spl4_7 | ~ spl4_99 | ~ spl4_120), inference(avatar_split_clause, [], [f10739, f1111, f1023, f637])).
fof(f10739, plain, ((e1 = op(e4, e3)) | (~ spl4_99 | ~ spl4_120)), inference(backward_demodulation, [], [f10645, f1025])).
fof(f10645, plain, ((e1 = op(e4, op(e1, e0))) | ~ spl4_120), inference(forward_demodulation, [], [f336, f1113])).
fof(f10741, plain, (spl4_26 | ~ spl4_99 | ~ spl4_120), inference(avatar_split_clause, [], [f10738, f1111, f1023, f717])).
fof(f10738, plain, ((e0 = op(e3, e4)) | (~ spl4_99 | ~ spl4_120)), inference(backward_demodulation, [], [f10644, f1025])).
fof(f10644, plain, ((e0 = op(op(e1, e0), e4)) | ~ spl4_120), inference(forward_demodulation, [], [f332, f1113])).
fof(f10740, plain, (~ spl4_89 | ~ spl4_99), inference(avatar_split_clause, [], [f10734, f1023, f981])).
fof(f981, plain, (spl4_89 <=> (e3 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl4_89])])).
fof(f10734, plain, (~ (e3 = op(e1, e2)) | ~ spl4_99), inference(backward_demodulation, [], [f417, f1025])).
fof(f10718, plain, (~ spl4_89 | ~ spl4_114), inference(avatar_split_clause, [], [f10712, f1086, f981])).
fof(f10712, plain, (~ (e3 = op(e1, e2)) | ~ spl4_114), inference(backward_demodulation, [], [f376, f1088])).
fof(f10697, plain, (~ spl4_120 | ~ spl4_122 | spl4_166), inference(avatar_contradiction_clause, [], [f10696])).
fof(f10696, plain, ($false | (~ spl4_120 | ~ spl4_122 | spl4_166)), inference(subsumption_resolution, [], [f10685, f1113])).
fof(f10685, plain, (~ (e4 = op(e0, e1)) | (~ spl4_122 | spl4_166)), inference(backward_demodulation, [], [f1406, f1122])).
fof(f1122, plain, ((op(e0, e0) = e1) | ~ spl4_122), inference(avatar_component_clause, [], [f1120])).
fof(f1120, plain, (spl4_122 <=> (op(e0, e0) = e1)), introduced(avatar_definition, [new_symbols(naming, [spl4_122])])).
fof(f1406, plain, (~ (e4 = op(e0, op(e0, e0))) | spl4_166), inference(avatar_component_clause, [], [f1404])).
fof(f1404, plain, (spl4_166 <=> (e4 = op(e0, op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl4_166])])).
fof(f10695, plain, (~ spl4_107 | ~ spl4_122), inference(avatar_split_clause, [], [f10681, f1120, f1057])).
fof(f10681, plain, (~ (e1 = op(e0, e3)) | ~ spl4_122), inference(backward_demodulation, [], [f408, f1122])).
fof(f10694, plain, (~ spl4_97 | ~ spl4_122), inference(avatar_split_clause, [], [f10677, f1120, f1015])).
fof(f10677, plain, (~ (e1 = op(e1, e0)) | ~ spl4_122), inference(backward_demodulation, [], [f356, f1122])).
fof(f10693, plain, (spl4_91 | ~ spl4_122), inference(avatar_split_clause, [], [f10676, f1120, f990])).
fof(f10676, plain, ((e0 = op(e1, e1)) | ~ spl4_122), inference(backward_demodulation, [], [f331, f1122])).
fof(f10673, plain, (~ spl4_18 | spl4_143 | ~ spl4_147), inference(avatar_split_clause, [], [f10672, f1305, f1287, f683])).
fof(f10672, plain, (~ (e2 = op(e4, e1)) | (spl4_143 | ~ spl4_147)), inference(forward_demodulation, [], [f1289, f1306])).
fof(f1306, plain, ((e1 = op(e4, op(e4, e4))) | ~ spl4_147), inference(avatar_component_clause, [], [f1305])).
fof(f10665, plain, (~ spl4_122 | spl4_158 | ~ spl4_169), inference(avatar_split_clause, [], [f10661, f1419, f1364, f1120])).
fof(f1364, plain, (spl4_158 <=> (e1 = op(op(e3, op(e3, e3)), op(e3, op(e3, e3))))), introduced(avatar_definition, [new_symbols(naming, [spl4_158])])).
fof(f10661, plain, (~ (op(e0, e0) = e1) | (spl4_158 | ~ spl4_169)), inference(backward_demodulation, [], [f1366, f1420])).
fof(f1366, plain, (~ (e1 = op(op(e3, op(e3, e3)), op(e3, op(e3, e3)))) | spl4_158), inference(avatar_component_clause, [], [f1364])).
fof(f10641, plain, (~ spl4_94 | ~ spl4_44), inference(avatar_split_clause, [], [f10640, f792, f1002])).
fof(f10640, plain, (~ (e3 = op(e1, e1)) | ~ spl4_44), inference(forward_demodulation, [], [f371, f794])).
fof(f10639, plain, (~ spl4_95 | ~ spl4_120), inference(avatar_split_clause, [], [f10638, f1111, f1006])).
fof(f1006, plain, (spl4_95 <=> (e4 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl4_95])])).
fof(f10638, plain, (~ (e4 = op(e1, e1)) | ~ spl4_120), inference(forward_demodulation, [], [f366, f1113])).
fof(f10626, plain, (spl4_1 | ~ spl4_5 | ~ spl4_133), inference(avatar_split_clause, [], [f10625, f1241, f628, f612])).
fof(f1241, plain, (spl4_133 <=> (e0 = op(op(e4, op(e4, e4)), op(e4, op(e4, e4))))), introduced(avatar_definition, [new_symbols(naming, [spl4_133])])).
fof(f10625, plain, ((e0 = op(e4, e4)) | (~ spl4_5 | ~ spl4_133)), inference(forward_demodulation, [], [f10620, f630])).
fof(f10620, plain, ((e0 = op(op(e4, e4), op(e4, e4))) | (~ spl4_5 | ~ spl4_133)), inference(backward_demodulation, [], [f1242, f630])).
fof(f1242, plain, ((e0 = op(op(e4, op(e4, e4)), op(e4, op(e4, e4)))) | ~ spl4_133), inference(avatar_component_clause, [], [f1241])).
fof(f10599, plain, (~ spl4_1 | ~ spl4_21), inference(avatar_split_clause, [], [f10592, f696, f612])).
fof(f10592, plain, (~ (e0 = op(e4, e4)) | ~ spl4_21), inference(backward_demodulation, [], [f449, f698])).
fof(f10579, plain, (~ spl4_31 | ~ spl4_48 | ~ spl4_62 | spl4_158), inference(avatar_contradiction_clause, [], [f10578])).
fof(f10578, plain, ($false | (~ spl4_31 | ~ spl4_48 | ~ spl4_62 | spl4_158)), inference(subsumption_resolution, [], [f10577, f870])).
fof(f10577, plain, (~ (e1 = op(e2, e2)) | (~ spl4_31 | ~ spl4_48 | spl4_158)), inference(forward_demodulation, [], [f10566, f811])).
fof(f811, plain, ((e2 = op(e3, e0)) | ~ spl4_48), inference(avatar_component_clause, [], [f809])).
fof(f10566, plain, (~ (e1 = op(op(e3, e0), op(e3, e0))) | (~ spl4_31 | spl4_158)), inference(backward_demodulation, [], [f1366, f740])).
fof(f10576, plain, (~ spl4_31 | ~ spl4_40 | ~ spl4_48 | spl4_156), inference(avatar_contradiction_clause, [], [f10575])).
fof(f10575, plain, ($false | (~ spl4_31 | ~ spl4_40 | ~ spl4_48 | spl4_156)), inference(subsumption_resolution, [], [f10574, f777])).
fof(f10574, plain, (~ (e4 = op(e3, e2)) | (~ spl4_31 | ~ spl4_48 | spl4_156)), inference(forward_demodulation, [], [f10565, f811])).
fof(f10565, plain, (~ (e4 = op(e3, op(e3, e0))) | (~ spl4_31 | spl4_156)), inference(backward_demodulation, [], [f1353, f740])).
fof(f1353, plain, (~ (e4 = op(e3, op(e3, op(e3, e3)))) | spl4_156), inference(avatar_component_clause, [], [f1351])).
fof(f1351, plain, (spl4_156 <=> (e4 = op(e3, op(e3, op(e3, e3))))), introduced(avatar_definition, [new_symbols(naming, [spl4_156])])).
fof(f10573, plain, (~ spl4_31 | ~ spl4_48 | spl4_141), inference(avatar_contradiction_clause, [], [f10572])).
fof(f10572, plain, ($false | (~ spl4_31 | ~ spl4_48 | spl4_141)), inference(subsumption_resolution, [], [f10563, f811])).
fof(f10563, plain, (~ (e2 = op(e3, e0)) | (~ spl4_31 | spl4_141)), inference(backward_demodulation, [], [f1279, f740])).
fof(f1279, plain, (~ (e2 = op(e3, op(e3, e3))) | spl4_141), inference(avatar_component_clause, [], [f1277])).
fof(f1277, plain, (spl4_141 <=> (e2 = op(e3, op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl4_141])])).
fof(f10546, plain, (spl4_9 | ~ spl4_44 | ~ spl4_85), inference(avatar_split_clause, [], [f10543, f964, f792, f645])).
fof(f10543, plain, ((e3 = op(e4, e3)) | (~ spl4_44 | ~ spl4_85)), inference(backward_demodulation, [], [f10489, f794])).
fof(f10489, plain, ((e3 = op(e4, op(e3, e1))) | ~ spl4_85), inference(backward_demodulation, [], [f347, f966])).
fof(f10545, plain, (spl4_27 | ~ spl4_44 | ~ spl4_85), inference(avatar_split_clause, [], [f10542, f964, f792, f721])).
fof(f10542, plain, ((e1 = op(e3, e4)) | (~ spl4_44 | ~ spl4_85)), inference(backward_demodulation, [], [f10488, f794])).
fof(f10488, plain, ((e1 = op(op(e3, e1), e4)) | ~ spl4_85), inference(backward_demodulation, [], [f339, f966])).
fof(f10544, plain, (~ spl4_34 | ~ spl4_44), inference(avatar_split_clause, [], [f10539, f792, f750])).
fof(f10539, plain, (~ (e3 = op(e3, e3)) | ~ spl4_44), inference(backward_demodulation, [], [f441, f794])).
fof(f10536, plain, (~ spl4_38 | ~ spl4_48), inference(avatar_split_clause, [], [f10528, f809, f767])).
fof(f10528, plain, (~ (e2 = op(e3, e2)) | ~ spl4_48), inference(backward_demodulation, [], [f437, f811])).
fof(f10535, plain, (~ spl4_23 | ~ spl4_48), inference(avatar_split_clause, [], [f10526, f809, f704])).
fof(f10526, plain, (~ (e2 = op(e4, e0)) | ~ spl4_48), inference(backward_demodulation, [], [f365, f811])).
fof(f365, plain, ~ (op(e3, e0) = op(e4, e0)), inference(cnf_transformation, [], [f4])).
fof(f10520, plain, (~ spl4_53 | ~ spl4_58), inference(avatar_split_clause, [], [f10517, f851, f830])).
fof(f10517, plain, (~ (e2 = op(e2, e4)) | ~ spl4_58), inference(backward_demodulation, [], [f435, f853])).
fof(f10512, plain, (spl4_21 | ~ spl4_75 | ~ spl4_111), inference(avatar_split_clause, [], [f10509, f1074, f922, f696])).
fof(f922, plain, (spl4_75 <=> (e4 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl4_75])])).
fof(f10509, plain, ((e0 = op(e4, e0)) | (~ spl4_75 | ~ spl4_111)), inference(backward_demodulation, [], [f10201, f924])).
fof(f924, plain, ((e4 = op(e2, e0)) | ~ spl4_75), inference(avatar_component_clause, [], [f922])).
fof(f10201, plain, ((e0 = op(op(e2, e0), e0)) | ~ spl4_111), inference(backward_demodulation, [], [f333, f1076])).
fof(f10511, plain, (~ spl4_60 | ~ spl4_75), inference(avatar_split_clause, [], [f10507, f922, f859])).
fof(f10507, plain, (~ (e4 = op(e2, e3)) | ~ spl4_75), inference(backward_demodulation, [], [f428, f924])).
fof(f10504, plain, (~ spl4_1 | ~ spl4_76), inference(avatar_split_clause, [], [f10499, f927, f612])).
fof(f10499, plain, (~ (e0 = op(e4, e4)) | ~ spl4_76), inference(backward_demodulation, [], [f402, f929])).
fof(f10494, plain, (~ spl4_60 | ~ spl4_85), inference(avatar_split_clause, [], [f10490, f964, f859])).
fof(f10490, plain, (~ (e4 = op(e2, e3)) | ~ spl4_85), inference(backward_demodulation, [], [f390, f966])).
fof(f10485, plain, (spl4_48 | ~ spl4_66 | ~ spl4_89), inference(avatar_split_clause, [], [f10479, f981, f885, f809])).
fof(f10479, plain, ((e2 = op(e3, e0)) | (~ spl4_66 | ~ spl4_89)), inference(backward_demodulation, [], [f10405, f983])).
fof(f983, plain, ((e3 = op(e1, e2)) | ~ spl4_89), inference(avatar_component_clause, [], [f981])).
fof(f10405, plain, ((e2 = op(op(e1, e2), e0)) | ~ spl4_66), inference(forward_demodulation, [], [f342, f887])).
fof(f10484, plain, (~ spl4_79 | ~ spl4_89), inference(avatar_split_clause, [], [f10477, f981, f939])).
fof(f10477, plain, (~ (e3 = op(e1, e4)) | ~ spl4_89), inference(backward_demodulation, [], [f424, f983])).
fof(f10483, plain, (~ spl4_14 | ~ spl4_89), inference(avatar_split_clause, [], [f10475, f981, f666])).
fof(f666, plain, (spl4_14 <=> (e3 = op(e4, e2))), introduced(avatar_definition, [new_symbols(naming, [spl4_14])])).
fof(f10475, plain, (~ (e3 = op(e4, e2)) | ~ spl4_89), inference(backward_demodulation, [], [f382, f983])).
fof(f10473, plain, (~ spl4_53 | ~ spl4_103), inference(avatar_split_clause, [], [f10468, f1040, f830])).
fof(f10468, plain, (~ (e2 = op(e2, e4)) | ~ spl4_103), inference(backward_demodulation, [], [f397, f1042])).
fof(f10464, plain, (~ spl4_102 | ~ spl4_107), inference(avatar_split_clause, [], [f10459, f1057, f1036])).
fof(f10459, plain, (~ (e1 = op(e0, e4)) | ~ spl4_107), inference(backward_demodulation, [], [f415, f1059])).
fof(f10463, plain, (~ spl4_7 | ~ spl4_107), inference(avatar_split_clause, [], [f10458, f1057, f637])).
fof(f10458, plain, (~ (e1 = op(e4, e3)) | ~ spl4_107), inference(backward_demodulation, [], [f389, f1059])).
fof(f10451, plain, (spl4_17 | ~ spl4_97 | ~ spl4_120), inference(avatar_split_clause, [], [f10448, f1111, f1015, f679])).
fof(f10448, plain, ((e1 = op(e4, e1)) | (~ spl4_97 | ~ spl4_120)), inference(backward_demodulation, [], [f10410, f1113])).
fof(f10410, plain, ((e1 = op(op(e0, e1), e1)) | ~ spl4_97), inference(forward_demodulation, [], [f336, f1017])).
fof(f10450, plain, (~ spl4_20 | ~ spl4_120), inference(avatar_split_clause, [], [f10445, f1111, f691])).
fof(f691, plain, (spl4_20 <=> (e4 = op(e4, e1))), introduced(avatar_definition, [new_symbols(naming, [spl4_20])])).
fof(f10445, plain, (~ (e4 = op(e4, e1)) | ~ spl4_120), inference(backward_demodulation, [], [f369, f1113])).
fof(f10443, plain, (spl4_107 | ~ spl4_124 | ~ spl4_180), inference(avatar_split_clause, [], [f10439, f1492, f1128, f1057])).
fof(f10439, plain, ((e1 = op(e0, e3)) | (~ spl4_124 | ~ spl4_180)), inference(backward_demodulation, [], [f1493, f1130])).
fof(f10442, plain, (~ spl4_119 | ~ spl4_124), inference(avatar_split_clause, [], [f10435, f1128, f1107])).
fof(f1107, plain, (spl4_119 <=> (e3 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl4_119])])).
fof(f10435, plain, (~ (e3 = op(e0, e1)) | ~ spl4_124), inference(backward_demodulation, [], [f406, f1130])).
fof(f406, plain, ~ (op(e0, e0) = op(e0, e1)), inference(cnf_transformation, [], [f4])).
fof(f10441, plain, (~ spl4_74 | ~ spl4_124), inference(avatar_split_clause, [], [f10432, f1128, f918])).
fof(f10432, plain, (~ (e3 = op(e2, e0)) | ~ spl4_124), inference(backward_demodulation, [], [f357, f1130])).
fof(f357, plain, ~ (op(e0, e0) = op(e2, e0)), inference(cnf_transformation, [], [f4])).
fof(f10426, plain, (spl4_89 | ~ spl4_93 | ~ spl4_148), inference(avatar_split_clause, [], [f10425, f1310, f998, f981])).
fof(f1310, plain, (spl4_148 <=> (e3 = op(e1, op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl4_148])])).
fof(f10425, plain, ((e3 = op(e1, e2)) | (~ spl4_93 | ~ spl4_148)), inference(forward_demodulation, [], [f1311, f1000])).
fof(f1311, plain, ((e3 = op(e1, op(e1, e1))) | ~ spl4_148), inference(avatar_component_clause, [], [f1310])).
fof(f10417, plain, (~ spl4_119 | spl4_170 | ~ spl4_180), inference(avatar_split_clause, [], [f10416, f1492, f1426, f1107])).
fof(f10416, plain, (~ (e3 = op(e0, e1)) | (spl4_170 | ~ spl4_180)), inference(backward_demodulation, [], [f1428, f1493])).
fof(f10403, plain, (spl4_90 | ~ spl4_93 | ~ spl4_146), inference(avatar_split_clause, [], [f10377, f1300, f998, f985])).
fof(f10377, plain, ((e4 = op(e1, e2)) | (~ spl4_93 | ~ spl4_146)), inference(forward_demodulation, [], [f1301, f1000])).
fof(f1301, plain, ((e4 = op(e1, op(e1, e1))) | ~ spl4_146), inference(avatar_component_clause, [], [f1300])).
fof(f10392, plain, (~ spl4_43 | ~ spl4_93), inference(avatar_split_clause, [], [f10391, f998, f788])).
fof(f10391, plain, (~ (e2 = op(e3, e1)) | ~ spl4_93), inference(forward_demodulation, [], [f371, f1000])).
fof(f10384, plain, (~ spl4_1 | ~ spl4_23 | spl4_139), inference(avatar_contradiction_clause, [], [f10383])).
fof(f10383, plain, ($false | (~ spl4_1 | ~ spl4_23 | spl4_139)), inference(subsumption_resolution, [], [f10382, f706])).
fof(f10382, plain, (~ (e2 = op(e4, e0)) | (~ spl4_1 | spl4_139)), inference(forward_demodulation, [], [f1269, f614])).
fof(f1269, plain, (~ (e2 = op(e4, op(e4, e4))) | spl4_139), inference(avatar_component_clause, [], [f1267])).
fof(f1267, plain, (spl4_139 <=> (e2 = op(e4, op(e4, e4)))), introduced(avatar_definition, [new_symbols(naming, [spl4_139])])).
fof(f10374, plain, (~ spl4_102 | ~ spl4_119 | ~ spl4_125 | spl4_170), inference(avatar_contradiction_clause, [], [f10373])).
fof(f10373, plain, ($false | (~ spl4_102 | ~ spl4_119 | ~ spl4_125 | spl4_170)), inference(subsumption_resolution, [], [f10372, f1109])).
fof(f1109, plain, ((e3 = op(e0, e1)) | ~ spl4_119), inference(avatar_component_clause, [], [f1107])).
fof(f10372, plain, (~ (e3 = op(e0, e1)) | (~ spl4_102 | ~ spl4_125 | spl4_170)), inference(forward_demodulation, [], [f10371, f1038])).
fof(f10371, plain, (~ (e3 = op(e0, op(e0, e4))) | (~ spl4_125 | spl4_170)), inference(forward_demodulation, [], [f1428, f1134])).
fof(f10328, plain, (~ spl4_2 | ~ spl4_7), inference(avatar_split_clause, [], [f10322, f637, f616])).
fof(f10322, plain, (~ (e1 = op(e4, e4)) | ~ spl4_7), inference(backward_demodulation, [], [f455, f639])).
fof(f639, plain, ((e1 = op(e4, e3)) | ~ spl4_7), inference(avatar_component_clause, [], [f637])).
fof(f10321, plain, (spl4_7 | ~ spl4_20 | ~ spl4_79), inference(avatar_split_clause, [], [f10319, f939, f691, f637])).
fof(f10319, plain, ((e1 = op(e4, e3)) | (~ spl4_20 | ~ spl4_79)), inference(backward_demodulation, [], [f10280, f693])).
fof(f693, plain, ((e4 = op(e4, e1)) | ~ spl4_20), inference(avatar_component_clause, [], [f691])).
fof(f10280, plain, ((e1 = op(op(e4, e1), e3)) | ~ spl4_79), inference(backward_demodulation, [], [f340, f941])).
fof(f10316, plain, (~ spl4_18 | ~ spl4_23), inference(avatar_split_clause, [], [f10309, f704, f683])).
fof(f10309, plain, (~ (e2 = op(e4, e1)) | ~ spl4_23), inference(backward_demodulation, [], [f446, f706])).
fof(f446, plain, ~ (op(e4, e0) = op(e4, e1)), inference(cnf_transformation, [], [f4])).
fof(f10300, plain, (spl4_23 | ~ spl4_66 | ~ spl4_90), inference(avatar_split_clause, [], [f10298, f985, f885, f704])).
fof(f10298, plain, ((e2 = op(e4, e0)) | (~ spl4_66 | ~ spl4_90)), inference(backward_demodulation, [], [f10263, f887])).
fof(f10263, plain, ((e2 = op(e4, op(e2, e1))) | ~ spl4_90), inference(backward_demodulation, [], [f342, f987])).
fof(f10299, plain, (~ spl4_61 | ~ spl4_66), inference(avatar_split_clause, [], [f10292, f885, f864])).
fof(f10292, plain, (~ (e0 = op(e2, e2)) | ~ spl4_66), inference(backward_demodulation, [], [f430, f887])).
fof(f10279, plain, (~ spl4_76 | ~ spl4_81), inference(avatar_split_clause, [], [f10271, f948, f927])).
fof(f10271, plain, (~ (e0 = op(e1, e4)) | ~ spl4_81), inference(backward_demodulation, [], [f425, f950])).
fof(f10278, plain, (~ spl4_6 | ~ spl4_81), inference(avatar_split_clause, [], [f10270, f948, f633])).
fof(f633, plain, (spl4_6 <=> (e0 = op(e4, e3))), introduced(avatar_definition, [new_symbols(naming, [spl4_6])])).
fof(f10270, plain, (~ (e0 = op(e4, e3)) | ~ spl4_81), inference(backward_demodulation, [], [f392, f950])).
fof(f10242, plain, (spl4_81 | ~ spl4_97 | ~ spl4_119), inference(avatar_split_clause, [], [f10238, f1107, f1015, f948])).
fof(f10238, plain, ((e0 = op(e1, e3)) | (~ spl4_97 | ~ spl4_119)), inference(backward_demodulation, [], [f10190, f1017])).
fof(f10190, plain, ((e0 = op(op(e1, e0), e3)) | ~ spl4_119), inference(backward_demodulation, [], [f332, f1109])).
fof(f10241, plain, (~ spl4_87 | ~ spl4_97), inference(avatar_split_clause, [], [f10232, f1015, f973])).
fof(f10232, plain, (~ (e1 = op(e1, e2)) | ~ spl4_97), inference(backward_demodulation, [], [f417, f1017])).
fof(f10240, plain, (~ spl4_72 | ~ spl4_97), inference(avatar_split_clause, [], [f10229, f1015, f910])).
fof(f10229, plain, (~ (e1 = op(e2, e0)) | ~ spl4_97), inference(backward_demodulation, [], [f360, f1017])).
fof(f360, plain, ~ (op(e1, e0) = op(e2, e0)), inference(cnf_transformation, [], [f4])).
fof(f10228, plain, (~ spl4_2 | ~ spl4_102), inference(avatar_split_clause, [], [f10224, f1036, f616])).
fof(f10224, plain, (~ (e1 = op(e4, e4)) | ~ spl4_102), inference(backward_demodulation, [], [f399, f1038])).
fof(f10211, plain, (~ spl4_61 | ~ spl4_111), inference(avatar_split_clause, [], [f10204, f1074, f864])).
fof(f10204, plain, (~ (e0 = op(e2, e2)) | ~ spl4_111), inference(backward_demodulation, [], [f377, f1076])).
fof(f10200, plain, (~ spl4_104 | ~ spl4_119), inference(avatar_split_clause, [], [f10197, f1107, f1044])).
fof(f10197, plain, (~ (e3 = op(e0, e4)) | ~ spl4_119), inference(backward_demodulation, [], [f412, f1109])).
fof(f10199, plain, (~ spl4_69 | ~ spl4_119), inference(avatar_split_clause, [], [f10193, f1107, f897])).
fof(f10193, plain, (~ (e3 = op(e2, e1)) | ~ spl4_119), inference(backward_demodulation, [], [f367, f1109])).
fof(f10188, plain, (~ spl4_115 | ~ spl4_125), inference(avatar_split_clause, [], [f10182, f1132, f1090])).
fof(f10182, plain, (~ (e4 = op(e0, e2)) | ~ spl4_125), inference(backward_demodulation, [], [f407, f1134])).
fof(f10187, plain, (~ spl4_25 | ~ spl4_125), inference(avatar_split_clause, [], [f10180, f1132, f712])).
fof(f10180, plain, (~ (e4 = op(e4, e0)) | ~ spl4_125), inference(backward_demodulation, [], [f359, f1134])).
fof(f10140, plain, (~ spl4_64 | ~ spl4_139 | spl4_183), inference(avatar_split_clause, [], [f10139, f1513, f1267, f876])).
fof(f10139, plain, (~ (e3 = op(e2, e2)) | (~ spl4_139 | spl4_183)), inference(forward_demodulation, [], [f1515, f1268])).
fof(f1268, plain, ((e2 = op(e4, op(e4, e4))) | ~ spl4_139), inference(avatar_component_clause, [], [f1267])).
fof(f10138, plain, (~ spl4_4 | ~ spl4_14), inference(avatar_split_clause, [], [f10137, f666, f624])).
fof(f10137, plain, (~ (e3 = op(e4, e4)) | ~ spl4_14), inference(forward_demodulation, [], [f454, f668])).
fof(f668, plain, ((e3 = op(e4, e2)) | ~ spl4_14), inference(avatar_component_clause, [], [f666])).
fof(f10136, plain, (~ spl4_5 | ~ spl4_30), inference(avatar_split_clause, [], [f10135, f733, f628])).
fof(f10135, plain, (~ (e4 = op(e4, e4)) | ~ spl4_30), inference(forward_demodulation, [], [f405, f735])).
fof(f735, plain, ((e4 = op(e3, e4)) | ~ spl4_30), inference(avatar_component_clause, [], [f733])).
fof(f10134, plain, (~ spl4_3 | ~ spl4_53), inference(avatar_split_clause, [], [f10133, f830, f620])).
fof(f10133, plain, (~ (e2 = op(e4, e4)) | ~ spl4_53), inference(forward_demodulation, [], [f404, f832])).
fof(f832, plain, ((e2 = op(e2, e4)) | ~ spl4_53), inference(avatar_component_clause, [], [f830])).
fof(f10132, plain, (~ spl4_63 | ~ spl4_139 | spl4_175), inference(avatar_split_clause, [], [f10131, f1457, f1267, f872])).
fof(f10131, plain, (~ (e2 = op(e2, e2)) | (~ spl4_139 | spl4_175)), inference(forward_demodulation, [], [f1459, f1268])).
fof(f10122, plain, (~ spl4_2 | ~ spl4_14 | ~ spl4_18 | spl4_151), inference(avatar_contradiction_clause, [], [f10121])).
fof(f10121, plain, ($false | (~ spl4_2 | ~ spl4_14 | ~ spl4_18 | spl4_151)), inference(subsumption_resolution, [], [f10117, f668])).
fof(f10117, plain, (~ (e3 = op(e4, e2)) | (~ spl4_2 | ~ spl4_18 | spl4_151)), inference(backward_demodulation, [], [f9675, f685])).
fof(f9675, plain, (~ (e3 = op(e4, op(e4, e1))) | (~ spl4_2 | spl4_151)), inference(forward_demodulation, [], [f1327, f618])).
fof(f618, plain, ((e1 = op(e4, e4)) | ~ spl4_2), inference(avatar_component_clause, [], [f616])).
fof(f1327, plain, (~ (e3 = op(e4, op(e4, op(e4, e4)))) | spl4_151), inference(avatar_component_clause, [], [f1325])).
fof(f1325, plain, (spl4_151 <=> (e3 = op(e4, op(e4, op(e4, e4))))), introduced(avatar_definition, [new_symbols(naming, [spl4_151])])).
fof(f10112, plain, (~ spl4_10 | ~ spl4_25), inference(avatar_split_clause, [], [f10109, f712, f649])).
fof(f10109, plain, (~ (e4 = op(e4, e3)) | ~ spl4_25), inference(backward_demodulation, [], [f448, f714])).
fof(f10103, plain, (~ spl4_31 | ~ spl4_34 | spl4_169), inference(avatar_split_clause, [], [f10093, f1419, f750, f738])).
fof(f10093, plain, (~ (e0 = op(e3, e3)) | (~ spl4_34 | spl4_169)), inference(backward_demodulation, [], [f1421, f752])).
fof(f752, plain, ((e3 = op(e3, e3)) | ~ spl4_34), inference(avatar_component_clause, [], [f750])).
fof(f1421, plain, (~ (e0 = op(e3, op(e3, e3))) | spl4_169), inference(avatar_component_clause, [], [f1419])).
fof(f10075, plain, (~ spl4_37 | ~ spl4_42), inference(avatar_split_clause, [], [f10067, f784, f763])).
fof(f10067, plain, (~ (e1 = op(e3, e2)) | ~ spl4_42), inference(backward_demodulation, [], [f440, f786])).
fof(f440, plain, ~ (op(e3, e1) = op(e3, e2)), inference(cnf_transformation, [], [f4])).
fof(f10065, plain, (~ spl4_31 | ~ spl4_46), inference(avatar_split_clause, [], [f10057, f801, f738])).
fof(f10057, plain, (~ (e0 = op(e3, e3)) | ~ spl4_46), inference(backward_demodulation, [], [f438, f803])).
fof(f10052, plain, (~ spl4_28 | ~ spl4_53), inference(avatar_split_clause, [], [f10047, f830, f725])).
fof(f10047, plain, (~ (e2 = op(e3, e4)) | ~ spl4_53), inference(backward_demodulation, [], [f403, f832])).
fof(f10045, plain, (~ spl4_10 | ~ spl4_60), inference(avatar_split_clause, [], [f10043, f859, f649])).
fof(f10043, plain, (~ (e4 = op(e4, e3)) | ~ spl4_60), inference(backward_demodulation, [], [f394, f861])).
fof(f10030, plain, (spl4_42 | ~ spl4_69 | ~ spl4_87), inference(avatar_split_clause, [], [f10026, f973, f897, f784])).
fof(f10026, plain, ((e1 = op(e3, e1)) | (~ spl4_69 | ~ spl4_87)), inference(backward_demodulation, [], [f9985, f899])).
fof(f9985, plain, ((e1 = op(op(e2, e1), e1)) | ~ spl4_87), inference(backward_demodulation, [], [f338, f975])).
fof(f10019, plain, (spl4_18 | ~ spl4_72 | ~ spl4_115), inference(avatar_split_clause, [], [f10017, f1090, f910, f683])).
fof(f10017, plain, ((e2 = op(e4, e1)) | (~ spl4_72 | ~ spl4_115)), inference(backward_demodulation, [], [f9711, f912])).
fof(f9711, plain, ((e2 = op(e4, op(e2, e0))) | ~ spl4_115), inference(backward_demodulation, [], [f341, f1092])).
fof(f10018, plain, (~ spl4_67 | ~ spl4_72), inference(avatar_split_clause, [], [f10009, f910, f889])).
fof(f10009, plain, (~ (e1 = op(e2, e1)) | ~ spl4_72), inference(backward_demodulation, [], [f426, f912])).
fof(f9993, plain, (~ spl4_37 | ~ spl4_87), inference(avatar_split_clause, [], [f9988, f973, f763])).
fof(f9988, plain, (~ (e1 = op(e3, e2)) | ~ spl4_87), inference(backward_demodulation, [], [f381, f975])).
fof(f9983, plain, (~ spl4_79 | ~ spl4_99), inference(avatar_split_clause, [], [f9979, f1023, f939])).
fof(f9979, plain, (~ (e3 = op(e1, e4)) | ~ spl4_99), inference(backward_demodulation, [], [f419, f1025])).
fof(f9953, plain, (~ spl4_118 | ~ spl4_123), inference(avatar_split_clause, [], [f9946, f1124, f1103])).
fof(f9946, plain, (~ (e2 = op(e0, e1)) | ~ spl4_123), inference(backward_demodulation, [], [f406, f1126])).
fof(f9932, plain, (~ spl4_79 | ~ spl4_95 | spl4_148), inference(avatar_split_clause, [], [f9931, f1310, f1006, f939])).
fof(f9931, plain, (~ (e3 = op(e1, e4)) | (~ spl4_95 | spl4_148)), inference(forward_demodulation, [], [f1312, f1008])).
fof(f1008, plain, ((e4 = op(e1, e1)) | ~ spl4_95), inference(avatar_component_clause, [], [f1006])).
fof(f1312, plain, (~ (e3 = op(e1, op(e1, e1))) | spl4_148), inference(avatar_component_clause, [], [f1310])).
fof(f9922, plain, (~ spl4_16 | ~ spl4_2 | spl4_167), inference(avatar_split_clause, [], [f9921, f1409, f616, f675])).
fof(f9921, plain, (~ (e0 = op(e4, e1)) | (~ spl4_2 | spl4_167)), inference(forward_demodulation, [], [f1411, f618])).
fof(f1411, plain, (~ (e0 = op(e4, op(e4, e4))) | spl4_167), inference(avatar_component_clause, [], [f1409])).
fof(f9906, plain, (~ spl4_105 | ~ spl4_115), inference(avatar_split_clause, [], [f9905, f1090, f1048])).
fof(f9905, plain, (~ (e4 = op(e0, e4)) | ~ spl4_115), inference(forward_demodulation, [], [f414, f1092])).
fof(f9890, plain, (~ spl4_19 | ~ spl4_2 | spl4_135), inference(avatar_split_clause, [], [f9619, f1249, f616, f687])).
fof(f9619, plain, (~ (e3 = op(e4, e1)) | (~ spl4_2 | spl4_135)), inference(backward_demodulation, [], [f1251, f618])).
fof(f9886, plain, (spl4_10 | ~ spl4_14 | ~ spl4_55), inference(avatar_split_clause, [], [f9885, f838, f666, f649])).
fof(f9885, plain, ((e4 = op(e4, e3)) | (~ spl4_14 | ~ spl4_55)), inference(backward_demodulation, [], [f9820, f668])).
fof(f9820, plain, ((e4 = op(e4, op(e4, e2))) | ~ spl4_55), inference(backward_demodulation, [], [f353, f840])).
fof(f353, plain, (e4 = op(op(e2, e4), op(e4, e2))), inference(cnf_transformation, [], [f3])).
fof(f9867, plain, (~ spl4_2 | ~ spl4_31 | ~ spl4_50 | spl4_158), inference(avatar_contradiction_clause, [], [f9866])).
fof(f9866, plain, ($false | (~ spl4_2 | ~ spl4_31 | ~ spl4_50 | spl4_158)), inference(subsumption_resolution, [], [f9865, f618])).
fof(f9865, plain, (~ (e1 = op(e4, e4)) | (~ spl4_31 | ~ spl4_50 | spl4_158)), inference(forward_demodulation, [], [f9857, f819])).
fof(f9857, plain, (~ (e1 = op(op(e3, e0), op(e3, e0))) | (~ spl4_31 | spl4_158)), inference(backward_demodulation, [], [f1366, f740])).
fof(f9862, plain, (~ spl4_31 | ~ spl4_50 | spl4_132), inference(avatar_contradiction_clause, [], [f9861])).
fof(f9861, plain, ($false | (~ spl4_31 | ~ spl4_50 | spl4_132)), inference(subsumption_resolution, [], [f9854, f819])).
fof(f9854, plain, (~ (e4 = op(e3, e0)) | (~ spl4_31 | spl4_132)), inference(backward_demodulation, [], [f1238, f740])).
fof(f1238, plain, (~ (e4 = op(e3, op(e3, e3))) | spl4_132), inference(avatar_component_clause, [], [f1236])).
fof(f9830, plain, (~ spl4_30 | ~ spl4_50), inference(avatar_split_clause, [], [f9827, f817, f733])).
fof(f9827, plain, (~ (e4 = op(e3, e4)) | ~ spl4_50), inference(backward_demodulation, [], [f439, f819])).
fof(f9817, plain, (~ spl4_34 | ~ spl4_59), inference(avatar_split_clause, [], [f9813, f855, f750])).
fof(f9813, plain, (~ (e3 = op(e3, e3)) | ~ spl4_59), inference(backward_demodulation, [], [f393, f857])).
fof(f9803, plain, (~ spl4_38 | ~ spl4_63), inference(avatar_split_clause, [], [f9793, f872, f767])).
fof(f9793, plain, (~ (e2 = op(e3, e2)) | ~ spl4_63), inference(backward_demodulation, [], [f383, f874])).
fof(f9791, plain, (~ spl4_57 | ~ spl4_67), inference(avatar_split_clause, [], [f9783, f889, f847])).
fof(f9783, plain, (~ (e1 = op(e2, e3)) | ~ spl4_67), inference(backward_demodulation, [], [f431, f891])).
fof(f9790, plain, (~ spl4_42 | ~ spl4_67), inference(avatar_split_clause, [], [f9781, f889, f784])).
fof(f9781, plain, (~ (e1 = op(e3, e1)) | ~ spl4_67), inference(backward_demodulation, [], [f373, f891])).
fof(f9780, plain, (spl4_23 | ~ spl4_71 | ~ spl4_115), inference(avatar_split_clause, [], [f9777, f1090, f906, f704])).
fof(f906, plain, (spl4_71 <=> (e0 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl4_71])])).
fof(f9777, plain, ((e2 = op(e4, e0)) | (~ spl4_71 | ~ spl4_115)), inference(backward_demodulation, [], [f9711, f908])).
fof(f908, plain, ((e0 = op(e2, e0)) | ~ spl4_71), inference(avatar_component_clause, [], [f906])).
fof(f9779, plain, (~ spl4_61 | ~ spl4_71), inference(avatar_split_clause, [], [f9769, f906, f864])).
fof(f9769, plain, (~ (e0 = op(e2, e2)) | ~ spl4_71), inference(backward_demodulation, [], [f427, f908])).
fof(f9763, plain, (~ spl4_79 | ~ spl4_83 | ~ spl4_95 | spl4_145), inference(avatar_contradiction_clause, [], [f9762])).
fof(f9762, plain, ($false | (~ spl4_79 | ~ spl4_83 | ~ spl4_95 | spl4_145)), inference(subsumption_resolution, [], [f9757, f958])).
fof(f9757, plain, (~ (e2 = op(e1, e3)) | (~ spl4_79 | ~ spl4_95 | spl4_145)), inference(backward_demodulation, [], [f9665, f941])).
fof(f9665, plain, (~ (e2 = op(e1, op(e1, e4))) | (~ spl4_95 | spl4_145)), inference(forward_demodulation, [], [f1298, f1008])).
fof(f1298, plain, (~ (e2 = op(e1, op(e1, op(e1, e1)))) | spl4_145), inference(avatar_component_clause, [], [f1296])).
fof(f1296, plain, (spl4_145 <=> (e2 = op(e1, op(e1, op(e1, e1))))), introduced(avatar_definition, [new_symbols(naming, [spl4_145])])).
fof(f9761, plain, (spl4_50 | ~ spl4_16 | ~ spl4_79), inference(avatar_split_clause, [], [f9755, f939, f675, f817])).
fof(f9755, plain, ((e4 = op(e3, e0)) | (~ spl4_16 | ~ spl4_79)), inference(backward_demodulation, [], [f9654, f941])).
fof(f9654, plain, ((e4 = op(op(e1, e4), e0)) | ~ spl4_16), inference(forward_demodulation, [], [f352, f677])).
fof(f677, plain, ((e0 = op(e4, e1)) | ~ spl4_16), inference(avatar_component_clause, [], [f675])).
fof(f9760, plain, (~ spl4_54 | ~ spl4_79), inference(avatar_split_clause, [], [f9752, f939, f834])).
fof(f9752, plain, (~ (e3 = op(e2, e4)) | ~ spl4_79), inference(backward_demodulation, [], [f400, f941])).
fof(f9751, plain, (~ spl4_78 | ~ spl4_83), inference(avatar_split_clause, [], [f9747, f956, f935])).
fof(f9747, plain, (~ (e2 = op(e1, e4)) | ~ spl4_83), inference(backward_demodulation, [], [f425, f958])).
fof(f9750, plain, (~ spl4_8 | ~ spl4_83), inference(avatar_split_clause, [], [f9746, f956, f641])).
fof(f9746, plain, (~ (e2 = op(e4, e3)) | ~ spl4_83), inference(backward_demodulation, [], [f392, f958])).
fof(f9741, plain, (~ spl4_81 | ~ spl4_86), inference(avatar_split_clause, [], [f9734, f969, f948])).
fof(f9734, plain, (~ (e0 = op(e1, e3)) | ~ spl4_86), inference(backward_demodulation, [], [f423, f971])).
fof(f9728, plain, (~ spl4_57 | ~ spl4_107), inference(avatar_split_clause, [], [f9722, f1057, f847])).
fof(f9722, plain, (~ (e1 = op(e2, e3)) | ~ spl4_107), inference(backward_demodulation, [], [f387, f1059])).
fof(f387, plain, ~ (op(e0, e3) = op(e2, e3)), inference(cnf_transformation, [], [f4])).
fof(f9718, plain, (~ spl4_110 | ~ spl4_115), inference(avatar_split_clause, [], [f9716, f1090, f1069])).
fof(f9716, plain, (~ (e4 = op(e0, e3)) | ~ spl4_115), inference(backward_demodulation, [], [f413, f1092])).
fof(f9717, plain, (~ spl4_15 | ~ spl4_115), inference(avatar_split_clause, [], [f9715, f1090, f670])).
fof(f9715, plain, (~ (e4 = op(e4, e2)) | ~ spl4_115), inference(backward_demodulation, [], [f379, f1092])).
fof(f9709, plain, (spl4_86 | ~ spl4_97 | ~ spl4_118), inference(avatar_split_clause, [], [f9706, f1103, f1015, f969])).
fof(f9706, plain, ((e0 = op(e1, e2)) | (~ spl4_97 | ~ spl4_118)), inference(backward_demodulation, [], [f9661, f1105])).
fof(f9661, plain, ((e0 = op(e1, op(e0, e1))) | ~ spl4_97), inference(forward_demodulation, [], [f332, f1017])).
fof(f9708, plain, (spl4_67 | ~ spl4_97 | ~ spl4_118), inference(avatar_split_clause, [], [f9705, f1103, f1015, f889])).
fof(f9705, plain, ((e1 = op(e2, e1)) | (~ spl4_97 | ~ spl4_118)), inference(backward_demodulation, [], [f9660, f1105])).
fof(f9660, plain, ((e1 = op(op(e0, e1), e1)) | ~ spl4_97), inference(forward_demodulation, [], [f336, f1017])).
fof(f9696, plain, (~ spl4_24 | ~ spl4_124), inference(avatar_split_clause, [], [f9689, f1128, f708])).
fof(f9689, plain, (~ (e3 = op(e4, e0)) | ~ spl4_124), inference(backward_demodulation, [], [f359, f1130])).
fof(f9677, plain, (~ spl4_24 | ~ spl4_2 | ~ spl4_16 | spl4_151), inference(avatar_split_clause, [], [f9676, f1325, f675, f616, f708])).
fof(f9676, plain, (~ (e3 = op(e4, e0)) | (~ spl4_2 | ~ spl4_16 | spl4_151)), inference(forward_demodulation, [], [f9675, f677])).
fof(f9674, plain, (~ spl4_78 | ~ spl4_95 | spl4_153), inference(avatar_split_clause, [], [f9673, f1336, f1006, f935])).
fof(f9673, plain, (~ (e2 = op(e1, e4)) | (~ spl4_95 | spl4_153)), inference(forward_demodulation, [], [f1338, f1008])).
fof(f9662, plain, (~ spl4_124 | ~ spl4_2 | ~ spl4_16 | spl4_183), inference(avatar_split_clause, [], [f9637, f1513, f675, f616, f1128])).
fof(f9637, plain, (~ (op(e0, e0) = e3) | (~ spl4_2 | ~ spl4_16 | spl4_183)), inference(forward_demodulation, [], [f9627, f677])).
fof(f9627, plain, (~ (e3 = op(op(e4, e1), op(e4, e1))) | (~ spl4_2 | spl4_183)), inference(backward_demodulation, [], [f1515, f618])).
fof(f9646, plain, (~ spl4_45 | ~ spl4_95), inference(avatar_split_clause, [], [f9645, f1006, f796])).
fof(f9645, plain, (~ (e4 = op(e3, e1)) | ~ spl4_95), inference(forward_demodulation, [], [f371, f1008])).
fof(f9603, plain, (~ spl4_11 | ~ spl4_16), inference(avatar_split_clause, [], [f9594, f675, f654])).
fof(f9594, plain, (~ (e0 = op(e4, e2)) | ~ spl4_16), inference(backward_demodulation, [], [f450, f677])).
fof(f9589, plain, (~ spl4_52 | ~ spl4_57), inference(avatar_split_clause, [], [f9582, f847, f826])).
fof(f9582, plain, (~ (e1 = op(e2, e4)) | ~ spl4_57), inference(backward_demodulation, [], [f435, f849])).
fof(f849, plain, ((e1 = op(e2, e3)) | ~ spl4_57), inference(avatar_component_clause, [], [f847])).
fof(f9588, plain, (~ spl4_7 | ~ spl4_57), inference(avatar_split_clause, [], [f9581, f847, f637])).
fof(f9581, plain, (~ (e1 = op(e4, e3)) | ~ spl4_57), inference(backward_demodulation, [], [f394, f849])).
fof(f9575, plain, (~ spl4_56 | ~ spl4_61), inference(avatar_split_clause, [], [f9563, f864, f843])).
fof(f9563, plain, (~ (e0 = op(e2, e3)) | ~ spl4_61), inference(backward_demodulation, [], [f433, f866])).
fof(f9560, plain, (spl4_57 | ~ spl4_68 | ~ spl4_89), inference(avatar_split_clause, [], [f9558, f981, f893, f847])).
fof(f9558, plain, ((e1 = op(e2, e3)) | (~ spl4_68 | ~ spl4_89)), inference(backward_demodulation, [], [f9519, f895])).
fof(f9519, plain, ((e1 = op(op(e2, e1), e3)) | ~ spl4_89), inference(backward_demodulation, [], [f338, f983])).
fof(f9551, plain, (spl4_16 | ~ spl4_75 | ~ spl4_112), inference(avatar_split_clause, [], [f9548, f1078, f922, f675])).
fof(f9548, plain, ((e0 = op(e4, e1)) | (~ spl4_75 | ~ spl4_112)), inference(backward_demodulation, [], [f9487, f924])).
fof(f9487, plain, ((e0 = op(op(e2, e0), e1)) | ~ spl4_112), inference(backward_demodulation, [], [f333, f1080])).
fof(f9550, plain, (~ spl4_65 | ~ spl4_75), inference(avatar_split_clause, [], [f9545, f922, f880])).
fof(f9545, plain, (~ (e4 = op(e2, e2)) | ~ spl4_75), inference(backward_demodulation, [], [f427, f924])).
fof(f9536, plain, (~ spl4_56 | ~ spl4_81), inference(avatar_split_clause, [], [f9527, f948, f843])).
fof(f9527, plain, (~ (e0 = op(e2, e3)) | ~ spl4_81), inference(backward_demodulation, [], [f390, f950])).
fof(f9518, plain, (spl4_78 | ~ spl4_95 | ~ spl4_153), inference(avatar_split_clause, [], [f9514, f1336, f1006, f935])).
fof(f9514, plain, ((e2 = op(e1, e4)) | (~ spl4_95 | ~ spl4_153)), inference(backward_demodulation, [], [f1337, f1008])).
fof(f1337, plain, ((e2 = op(e1, op(e1, e1))) | ~ spl4_153), inference(avatar_component_clause, [], [f1336])).
fof(f9517, plain, (~ spl4_20 | ~ spl4_95), inference(avatar_split_clause, [], [f9510, f1006, f691])).
fof(f9510, plain, (~ (e4 = op(e4, e1)) | ~ spl4_95), inference(backward_demodulation, [], [f372, f1008])).
fof(f9516, plain, (spl4_2 | ~ spl4_95), inference(avatar_split_clause, [], [f9508, f1006, f616])).
fof(f9508, plain, ((e1 = op(e4, e4)) | ~ spl4_95), inference(backward_demodulation, [], [f337, f1008])).
fof(f9495, plain, (~ spl4_87 | ~ spl4_112), inference(avatar_split_clause, [], [f9489, f1078, f973])).
fof(f9489, plain, (~ (e1 = op(e1, e2)) | ~ spl4_112), inference(backward_demodulation, [], [f376, f1080])).
fof(f9486, plain, (~ spl4_114 | ~ spl4_119), inference(avatar_split_clause, [], [f9483, f1107, f1086])).
fof(f9483, plain, (~ (e3 = op(e0, e2)) | ~ spl4_119), inference(backward_demodulation, [], [f410, f1109])).
fof(f9459, plain, (~ spl4_61 | spl4_144 | ~ spl4_153), inference(avatar_split_clause, [], [f9456, f1336, f1292, f864])).
fof(f1292, plain, (spl4_144 <=> (e0 = op(op(e1, op(e1, e1)), op(e1, op(e1, e1))))), introduced(avatar_definition, [new_symbols(naming, [spl4_144])])).
fof(f9456, plain, (~ (e0 = op(e2, e2)) | (spl4_144 | ~ spl4_153)), inference(backward_demodulation, [], [f1294, f1337])).
fof(f1294, plain, (~ (e0 = op(op(e1, op(e1, e1)), op(e1, op(e1, e1)))) | spl4_144), inference(avatar_component_clause, [], [f1292])).
fof(f9441, plain, (spl4_31 | ~ spl4_34 | ~ spl4_130), inference(avatar_contradiction_clause, [], [f9440])).
fof(f9440, plain, ($false | (spl4_31 | ~ spl4_34 | ~ spl4_130)), inference(subsumption_resolution, [], [f9439, f739])).
fof(f739, plain, (~ (e0 = op(e3, e3)) | spl4_31), inference(avatar_component_clause, [], [f738])).
fof(f9439, plain, ((e0 = op(e3, e3)) | (~ spl4_34 | ~ spl4_130)), inference(forward_demodulation, [], [f9438, f752])).
fof(f9438, plain, ((e0 = op(op(e3, e3), op(e3, e3))) | (~ spl4_34 | ~ spl4_130)), inference(forward_demodulation, [], [f1229, f752])).
fof(f1229, plain, ((e0 = op(op(e3, op(e3, e3)), op(e3, op(e3, e3)))) | ~ spl4_130), inference(avatar_component_clause, [], [f1228])).
fof(f9385, plain, (~ spl4_5 | ~ spl4_20), inference(avatar_split_clause, [], [f9382, f691, f628])).
fof(f9382, plain, (~ (e4 = op(e4, e4)) | ~ spl4_20), inference(backward_demodulation, [], [f452, f693])).
fof(f452, plain, ~ (op(e4, e1) = op(e4, e4)), inference(cnf_transformation, [], [f4])).
fof(f9365, plain, (~ spl4_9 | ~ spl4_34), inference(avatar_split_clause, [], [f9360, f750, f645])).
fof(f9360, plain, (~ (e3 = op(e4, e3)) | ~ spl4_34), inference(backward_demodulation, [], [f395, f752])).
fof(f9358, plain, (~ spl4_33 | ~ spl4_38), inference(avatar_split_clause, [], [f9351, f767, f746])).
fof(f9351, plain, (~ (e2 = op(e3, e3)) | ~ spl4_38), inference(backward_demodulation, [], [f443, f769])).
fof(f443, plain, ~ (op(e3, e2) = op(e3, e3)), inference(cnf_transformation, [], [f4])).
fof(f9357, plain, (~ spl4_13 | ~ spl4_38), inference(avatar_split_clause, [], [f9350, f767, f662])).
fof(f9350, plain, (~ (e2 = op(e4, e2)) | ~ spl4_38), inference(backward_demodulation, [], [f385, f769])).
fof(f9348, plain, (~ spl4_17 | ~ spl4_42), inference(avatar_split_clause, [], [f9339, f784, f679])).
fof(f9339, plain, (~ (e1 = op(e4, e1)) | ~ spl4_42), inference(backward_demodulation, [], [f375, f786])).
fof(f375, plain, ~ (op(e3, e1) = op(e4, e1)), inference(cnf_transformation, [], [f4])).
fof(f9338, plain, (spl4_24 | ~ spl4_46 | ~ spl4_110), inference(avatar_split_clause, [], [f9335, f1069, f801, f708])).
fof(f9335, plain, ((e3 = op(e4, e0)) | (~ spl4_46 | ~ spl4_110)), inference(backward_demodulation, [], [f9005, f803])).
fof(f9005, plain, ((e3 = op(e4, op(e3, e0))) | ~ spl4_110), inference(backward_demodulation, [], [f346, f1071])).
fof(f9336, plain, (~ spl4_21 | ~ spl4_46), inference(avatar_split_clause, [], [f9325, f801, f696])).
fof(f9325, plain, (~ (e0 = op(e4, e0)) | ~ spl4_46), inference(backward_demodulation, [], [f365, f803])).
fof(f9298, plain, (~ spl4_64 | ~ spl4_69), inference(avatar_split_clause, [], [f9294, f897, f876])).
fof(f9294, plain, (~ (e3 = op(e2, e2)) | ~ spl4_69), inference(backward_demodulation, [], [f430, f899])).
fof(f9281, plain, (~ spl4_33 | ~ spl4_83), inference(avatar_split_clause, [], [f9277, f956, f746])).
fof(f9277, plain, (~ (e2 = op(e3, e3)) | ~ spl4_83), inference(backward_demodulation, [], [f391, f958])).
fof(f9274, plain, (~ spl4_82 | ~ spl4_87), inference(avatar_split_clause, [], [f9269, f973, f952])).
fof(f9269, plain, (~ (e1 = op(e1, e3)) | ~ spl4_87), inference(backward_demodulation, [], [f423, f975])).
fof(f9263, plain, (spl4_11 | ~ spl4_100 | ~ spl4_118), inference(avatar_split_clause, [], [f9255, f1103, f1027, f654])).
fof(f9255, plain, ((e0 = op(e4, e2)) | (~ spl4_100 | ~ spl4_118)), inference(backward_demodulation, [], [f9211, f1029])).
fof(f9211, plain, ((e0 = op(op(e1, e0), e2)) | ~ spl4_118), inference(backward_demodulation, [], [f332, f1105])).
fof(f9260, plain, (~ spl4_79 | ~ spl4_91 | ~ spl4_100 | spl4_152), inference(avatar_contradiction_clause, [], [f9259])).
fof(f9259, plain, ($false | (~ spl4_79 | ~ spl4_91 | ~ spl4_100 | spl4_152)), inference(subsumption_resolution, [], [f9251, f941])).
fof(f9251, plain, (~ (e3 = op(e1, e4)) | (~ spl4_91 | ~ spl4_100 | spl4_152)), inference(backward_demodulation, [], [f9043, f1029])).
fof(f9043, plain, (~ (e3 = op(e1, op(e1, e0))) | (~ spl4_91 | spl4_152)), inference(backward_demodulation, [], [f1332, f992])).
fof(f1332, plain, (~ (e3 = op(e1, op(e1, op(e1, e1)))) | spl4_152), inference(avatar_component_clause, [], [f1330])).
fof(f1330, plain, (spl4_152 <=> (e3 = op(e1, op(e1, op(e1, e1))))), introduced(avatar_definition, [new_symbols(naming, [spl4_152])])).
fof(f9257, plain, (~ spl4_75 | ~ spl4_100), inference(avatar_split_clause, [], [f9245, f1027, f922])).
fof(f9245, plain, (~ (e4 = op(e2, e0)) | ~ spl4_100), inference(backward_demodulation, [], [f360, f1029])).
fof(f9209, plain, (spl4_100 | ~ spl4_91 | ~ spl4_146), inference(avatar_split_clause, [], [f9208, f1300, f990, f1027])).
fof(f9208, plain, ((e4 = op(e1, e0)) | (~ spl4_91 | ~ spl4_146)), inference(forward_demodulation, [], [f1301, f992])).
fof(f9169, plain, (~ spl4_12 | ~ spl4_17), inference(avatar_split_clause, [], [f9163, f679, f658])).
fof(f9163, plain, (~ (e1 = op(e4, e2)) | ~ spl4_17), inference(backward_demodulation, [], [f450, f681])).
fof(f9162, plain, (~ spl4_16 | ~ spl4_21), inference(avatar_split_clause, [], [f9154, f696, f675])).
fof(f9154, plain, (~ (e0 = op(e4, e1)) | ~ spl4_21), inference(backward_demodulation, [], [f446, f698])).
fof(f9147, plain, (~ spl4_28 | ~ spl4_33), inference(avatar_split_clause, [], [f9142, f746, f725])).
fof(f9142, plain, (~ (e2 = op(e3, e4)) | ~ spl4_33), inference(backward_demodulation, [], [f445, f748])).
fof(f9140, plain, (~ spl4_12 | ~ spl4_37), inference(avatar_split_clause, [], [f9132, f763, f658])).
fof(f9132, plain, (~ (e1 = op(e4, e2)) | ~ spl4_37), inference(backward_demodulation, [], [f385, f765])).
fof(f765, plain, ((e1 = op(e3, e2)) | ~ spl4_37), inference(avatar_component_clause, [], [f763])).
fof(f9120, plain, (~ spl4_31 | ~ spl4_56), inference(avatar_split_clause, [], [f9110, f843, f738])).
fof(f9110, plain, (~ (e0 = op(e3, e3)) | ~ spl4_56), inference(backward_demodulation, [], [f393, f845])).
fof(f9107, plain, (~ spl4_39 | ~ spl4_64), inference(avatar_split_clause, [], [f9095, f876, f771])).
fof(f9095, plain, (~ (e3 = op(e3, e2)) | ~ spl4_64), inference(backward_demodulation, [], [f383, f878])).
fof(f9093, plain, (spl4_13 | ~ spl4_68 | ~ spl4_90), inference(avatar_split_clause, [], [f9091, f985, f893, f662])).
fof(f9091, plain, ((e2 = op(e4, e2)) | (~ spl4_68 | ~ spl4_90)), inference(backward_demodulation, [], [f9063, f895])).
fof(f9063, plain, ((e2 = op(e4, op(e2, e1))) | ~ spl4_90), inference(backward_demodulation, [], [f342, f987])).
fof(f9092, plain, (~ spl4_63 | ~ spl4_68), inference(avatar_split_clause, [], [f9086, f893, f872])).
fof(f9086, plain, (~ (e2 = op(e2, e2)) | ~ spl4_68), inference(backward_demodulation, [], [f430, f895])).
fof(f9061, plain, (~ spl4_64 | ~ spl4_91 | ~ spl4_98 | spl4_184), inference(avatar_split_clause, [], [f9060, f1518, f1019, f990, f876])).
fof(f1019, plain, (spl4_98 <=> (e2 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl4_98])])).
fof(f9060, plain, (~ (e3 = op(e2, e2)) | (~ spl4_91 | ~ spl4_98 | spl4_184)), inference(forward_demodulation, [], [f9047, f1021])).
fof(f1021, plain, ((e2 = op(e1, e0)) | ~ spl4_98), inference(avatar_component_clause, [], [f1019])).
fof(f9047, plain, (~ (e3 = op(op(e1, e0), op(e1, e0))) | (~ spl4_91 | spl4_184)), inference(backward_demodulation, [], [f1520, f992])).
fof(f9057, plain, (~ spl4_90 | ~ spl4_91 | ~ spl4_98 | spl4_157), inference(avatar_split_clause, [], [f9056, f1356, f1019, f990, f985])).
fof(f9056, plain, (~ (e4 = op(e1, e2)) | (~ spl4_91 | ~ spl4_98 | spl4_157)), inference(forward_demodulation, [], [f9045, f1021])).
fof(f9045, plain, (~ (e4 = op(e1, op(e1, e0))) | (~ spl4_91 | spl4_157)), inference(backward_demodulation, [], [f1358, f992])).
fof(f9055, plain, (~ spl4_91 | ~ spl4_98 | spl4_153), inference(avatar_contradiction_clause, [], [f9054])).
fof(f9054, plain, ($false | (~ spl4_91 | ~ spl4_98 | spl4_153)), inference(subsumption_resolution, [], [f9044, f1021])).
fof(f9044, plain, (~ (e2 = op(e1, e0)) | (~ spl4_91 | spl4_153)), inference(backward_demodulation, [], [f1338, f992])).
fof(f9049, plain, (~ spl4_86 | ~ spl4_91), inference(avatar_split_clause, [], [f9034, f990, f969])).
fof(f9034, plain, (~ (e0 = op(e1, e2)) | ~ spl4_91), inference(backward_demodulation, [], [f420, f992])).
fof(f420, plain, ~ (op(e1, e1) = op(e1, e2)), inference(cnf_transformation, [], [f4])).
fof(f9048, plain, (~ spl4_16 | ~ spl4_91), inference(avatar_split_clause, [], [f9033, f990, f675])).
fof(f9033, plain, (~ (e0 = op(e4, e1)) | ~ spl4_91), inference(backward_demodulation, [], [f372, f992])).
fof(f9030, plain, (spl4_37 | ~ spl4_98 | ~ spl4_119), inference(avatar_split_clause, [], [f9026, f1107, f1019, f763])).
fof(f9026, plain, ((e1 = op(e3, e2)) | (~ spl4_98 | ~ spl4_119)), inference(backward_demodulation, [], [f8978, f1021])).
fof(f8978, plain, ((e1 = op(e3, op(e1, e0))) | ~ spl4_119), inference(backward_demodulation, [], [f336, f1109])).
fof(f9029, plain, (spl4_56 | ~ spl4_98 | ~ spl4_119), inference(avatar_split_clause, [], [f9025, f1107, f1019, f843])).
fof(f9025, plain, ((e0 = op(e2, e3)) | (~ spl4_98 | ~ spl4_119)), inference(backward_demodulation, [], [f8977, f1021])).
fof(f8977, plain, ((e0 = op(op(e1, e0), e3)) | ~ spl4_119), inference(backward_demodulation, [], [f332, f1109])).
fof(f9028, plain, (~ spl4_83 | ~ spl4_98), inference(avatar_split_clause, [], [f9022, f1019, f956])).
fof(f9022, plain, (~ (e2 = op(e1, e3)) | ~ spl4_98), inference(backward_demodulation, [], [f418, f1021])).
fof(f9027, plain, (~ spl4_23 | ~ spl4_98), inference(avatar_split_clause, [], [f9019, f1019, f704])).
fof(f9019, plain, (~ (e2 = op(e4, e0)) | ~ spl4_98), inference(backward_demodulation, [], [f362, f1021])).
fof(f9010, plain, (~ spl4_60 | ~ spl4_110), inference(avatar_split_clause, [], [f9007, f1069, f859])).
fof(f9007, plain, (~ (e4 = op(e2, e3)) | ~ spl4_110), inference(backward_demodulation, [], [f387, f1071])).
fof(f9003, plain, (~ spl4_101 | ~ spl4_111), inference(avatar_split_clause, [], [f8997, f1074, f1032])).
fof(f8997, plain, (~ (e0 = op(e0, e4)) | ~ spl4_111), inference(backward_demodulation, [], [f414, f1076])).
fof(f8976, plain, (~ spl4_118 | ~ spl4_122 | spl4_171), inference(avatar_split_clause, [], [f8969, f1432, f1120, f1103])).
fof(f8969, plain, (~ (e2 = op(e0, e1)) | (~ spl4_122 | spl4_171)), inference(backward_demodulation, [], [f1434, f1122])).
fof(f8975, plain, (~ spl4_119 | ~ spl4_122 | spl4_168), inference(avatar_split_clause, [], [f8968, f1414, f1120, f1107])).
fof(f8968, plain, (~ (e3 = op(e0, e1)) | (~ spl4_122 | spl4_168)), inference(backward_demodulation, [], [f1416, f1122])).
fof(f1416, plain, (~ (e3 = op(e0, op(e0, e0))) | spl4_168), inference(avatar_component_clause, [], [f1414])).
fof(f8973, plain, (~ spl4_47 | ~ spl4_122), inference(avatar_split_clause, [], [f8957, f1120, f805])).
fof(f8957, plain, (~ (e1 = op(e3, e0)) | ~ spl4_122), inference(backward_demodulation, [], [f358, f1122])).
fof(f8907, plain, (~ spl4_40 | ~ spl4_45), inference(avatar_split_clause, [], [f8906, f796, f775])).
fof(f8906, plain, (~ (e4 = op(e3, e2)) | ~ spl4_45), inference(forward_demodulation, [], [f440, f798])).
fof(f8900, plain, (~ spl4_30 | ~ spl4_45), inference(avatar_split_clause, [], [f8899, f796, f733])).
fof(f8899, plain, (~ (e4 = op(e3, e4)) | ~ spl4_45), inference(forward_demodulation, [], [f442, f798])).
fof(f8890, plain, (spl4_1 | ~ spl4_5 | ~ spl4_167), inference(avatar_contradiction_clause, [], [f8889])).
fof(f8889, plain, ($false | (spl4_1 | ~ spl4_5 | ~ spl4_167)), inference(subsumption_resolution, [], [f8888, f613])).
fof(f613, plain, (~ (e0 = op(e4, e4)) | spl4_1), inference(avatar_component_clause, [], [f612])).
fof(f8888, plain, ((e0 = op(e4, e4)) | (~ spl4_5 | ~ spl4_167)), inference(backward_demodulation, [], [f1410, f630])).
fof(f8881, plain, (~ spl4_12 | ~ spl4_52 | ~ spl4_92), inference(avatar_contradiction_clause, [], [f8880])).
fof(f8880, plain, ($false | (~ spl4_12 | ~ spl4_52 | ~ spl4_92)), inference(subsumption_resolution, [], [f8879, f462])).
fof(f462, plain, ~ (e1 = e4), inference(cnf_transformation, [], [f5])).
fof(f8879, plain, ((e1 = e4) | (~ spl4_12 | ~ spl4_52 | ~ spl4_92)), inference(forward_demodulation, [], [f8875, f996])).
fof(f8875, plain, ((e4 = op(e1, e1)) | (~ spl4_12 | ~ spl4_52)), inference(backward_demodulation, [], [f8560, f660])).
fof(f8560, plain, ((e4 = op(e1, op(e4, e2))) | ~ spl4_52), inference(backward_demodulation, [], [f353, f828])).
fof(f828, plain, ((e1 = op(e2, e4)) | ~ spl4_52), inference(avatar_component_clause, [], [f826])).
fof(f8865, plain, (~ spl4_6 | ~ spl4_16), inference(avatar_split_clause, [], [f8857, f675, f633])).
fof(f8857, plain, (~ (e0 = op(e4, e3)) | ~ spl4_16), inference(backward_demodulation, [], [f451, f677])).
fof(f8841, plain, (~ spl4_6 | ~ spl4_31), inference(avatar_split_clause, [], [f8834, f738, f633])).
fof(f8834, plain, (~ (e0 = op(e4, e3)) | ~ spl4_31), inference(backward_demodulation, [], [f395, f740])).
fof(f8832, plain, (~ spl4_45 | ~ spl4_52 | ~ spl4_83), inference(avatar_contradiction_clause, [], [f8831])).
fof(f8831, plain, ($false | (~ spl4_45 | ~ spl4_52 | ~ spl4_83)), inference(subsumption_resolution, [], [f8830, f461])).
fof(f461, plain, ~ (e1 = e3), inference(cnf_transformation, [], [f5])).
fof(f8830, plain, ((e1 = e3) | (~ spl4_45 | ~ spl4_52 | ~ spl4_83)), inference(forward_demodulation, [], [f8827, f828])).
fof(f8827, plain, ((e3 = op(e2, e4)) | (~ spl4_45 | ~ spl4_83)), inference(backward_demodulation, [], [f8753, f798])).
fof(f8753, plain, ((e3 = op(e2, op(e3, e1))) | ~ spl4_83), inference(backward_demodulation, [], [f347, f958])).
fof(f8820, plain, (~ spl4_47 | spl4_91 | ~ spl4_107), inference(avatar_contradiction_clause, [], [f8819])).
fof(f8819, plain, ($false | (~ spl4_47 | spl4_91 | ~ spl4_107)), inference(subsumption_resolution, [], [f8816, f991])).
fof(f991, plain, (~ (e0 = op(e1, e1)) | spl4_91), inference(avatar_component_clause, [], [f990])).
fof(f8816, plain, ((e0 = op(e1, e1)) | (~ spl4_47 | ~ spl4_107)), inference(backward_demodulation, [], [f8499, f807])).
fof(f8499, plain, ((e0 = op(op(e3, e0), e1)) | ~ spl4_107), inference(backward_demodulation, [], [f334, f1059])).
fof(f8818, plain, (~ spl4_22 | ~ spl4_47), inference(avatar_split_clause, [], [f8810, f805, f700])).
fof(f8810, plain, (~ (e1 = op(e4, e0)) | ~ spl4_47), inference(backward_demodulation, [], [f365, f807])).
fof(f8718, plain, (~ spl4_18 | ~ spl4_118), inference(avatar_split_clause, [], [f8714, f1103, f683])).
fof(f8714, plain, (~ (e2 = op(e4, e1)) | ~ spl4_118), inference(backward_demodulation, [], [f369, f1105])).
fof(f8677, plain, (~ spl4_123 | ~ spl4_167 | spl4_175), inference(avatar_split_clause, [], [f8676, f1457, f1409, f1124])).
fof(f8676, plain, (~ (op(e0, e0) = e2) | (~ spl4_167 | spl4_175)), inference(forward_demodulation, [], [f1459, f1410])).
fof(f8647, plain, (~ spl4_62 | ~ spl4_141 | spl4_158), inference(avatar_split_clause, [], [f8646, f1364, f1277, f868])).
fof(f8646, plain, (~ (e1 = op(e2, e2)) | (~ spl4_141 | spl4_158)), inference(forward_demodulation, [], [f1366, f1278])).
fof(f1278, plain, ((e2 = op(e3, op(e3, e3))) | ~ spl4_141), inference(avatar_component_clause, [], [f1277])).
fof(f8633, plain, (~ spl4_4 | ~ spl4_6 | ~ spl4_22 | spl4_134), inference(avatar_contradiction_clause, [], [f8632])).
fof(f8632, plain, ($false | (~ spl4_4 | ~ spl4_6 | ~ spl4_22 | spl4_134)), inference(subsumption_resolution, [], [f8627, f702])).
fof(f8627, plain, (~ (e1 = op(e4, e0)) | (~ spl4_4 | ~ spl4_6 | spl4_134)), inference(backward_demodulation, [], [f8476, f635])).
fof(f635, plain, ((e0 = op(e4, e3)) | ~ spl4_6), inference(avatar_component_clause, [], [f633])).
fof(f8476, plain, (~ (e1 = op(e4, op(e4, e3))) | (~ spl4_4 | spl4_134)), inference(forward_demodulation, [], [f1247, f626])).
fof(f626, plain, ((e3 = op(e4, e4)) | ~ spl4_4), inference(avatar_component_clause, [], [f624])).
fof(f8615, plain, (spl4_15 | ~ spl4_18 | ~ spl4_80), inference(avatar_split_clause, [], [f8613, f943, f683, f670])).
fof(f8613, plain, ((e4 = op(e4, e2)) | (~ spl4_18 | ~ spl4_80)), inference(backward_demodulation, [], [f8318, f685])).
fof(f8318, plain, ((e4 = op(e4, op(e4, e1))) | ~ spl4_80), inference(backward_demodulation, [], [f352, f945])).
fof(f8614, plain, (~ spl4_8 | ~ spl4_18), inference(avatar_split_clause, [], [f8609, f683, f641])).
fof(f8609, plain, (~ (e2 = op(e4, e3)) | ~ spl4_18), inference(backward_demodulation, [], [f451, f685])).
fof(f8598, plain, (~ spl4_28 | ~ spl4_35 | ~ spl4_37 | spl4_131), inference(avatar_contradiction_clause, [], [f8597])).
fof(f8597, plain, ($false | (~ spl4_28 | ~ spl4_35 | ~ spl4_37 | spl4_131)), inference(subsumption_resolution, [], [f8593, f765])).
fof(f8593, plain, (~ (e1 = op(e3, e2)) | (~ spl4_28 | ~ spl4_35 | spl4_131)), inference(backward_demodulation, [], [f8477, f727])).
fof(f8477, plain, (~ (e1 = op(e3, op(e3, e4))) | (~ spl4_35 | spl4_131)), inference(forward_demodulation, [], [f1234, f756])).
fof(f8596, plain, (~ spl4_28 | ~ spl4_35 | ~ spl4_61 | spl4_130), inference(avatar_contradiction_clause, [], [f8595])).
fof(f8595, plain, ($false | (~ spl4_28 | ~ spl4_35 | ~ spl4_61 | spl4_130)), inference(subsumption_resolution, [], [f8591, f866])).
fof(f8591, plain, (~ (e0 = op(e2, e2)) | (~ spl4_28 | ~ spl4_35 | spl4_130)), inference(backward_demodulation, [], [f8191, f727])).
fof(f8191, plain, (~ (e0 = op(op(e3, e4), op(e3, e4))) | (~ spl4_35 | spl4_130)), inference(forward_demodulation, [], [f1230, f756])).
fof(f8585, plain, (~ spl4_37 | ~ spl4_39), inference(avatar_contradiction_clause, [], [f8584])).
fof(f8584, plain, ($false | (~ spl4_37 | ~ spl4_39)), inference(subsumption_resolution, [], [f8583, f461])).
fof(f8583, plain, ((e1 = e3) | (~ spl4_37 | ~ spl4_39)), inference(backward_demodulation, [], [f773, f765])).
fof(f8581, plain, (~ spl4_26 | ~ spl4_41), inference(avatar_split_clause, [], [f8573, f780, f717])).
fof(f8573, plain, (~ (e0 = op(e3, e4)) | ~ spl4_41), inference(backward_demodulation, [], [f442, f782])).
fof(f8580, plain, (~ spl4_16 | ~ spl4_41), inference(avatar_split_clause, [], [f8572, f780, f675])).
fof(f8572, plain, (~ (e0 = op(e4, e1)) | ~ spl4_41), inference(backward_demodulation, [], [f375, f782])).
fof(f8556, plain, (~ spl4_39 | ~ spl4_58), inference(avatar_contradiction_clause, [], [f8555])).
fof(f8555, plain, ($false | (~ spl4_39 | ~ spl4_58)), inference(subsumption_resolution, [], [f8554, f463])).
fof(f463, plain, ~ (e2 = e3), inference(cnf_transformation, [], [f5])).
fof(f8554, plain, ((e2 = e3) | (~ spl4_39 | ~ spl4_58)), inference(forward_demodulation, [], [f8550, f853])).
fof(f8550, plain, ((e3 = op(e2, e3)) | (~ spl4_39 | ~ spl4_58)), inference(backward_demodulation, [], [f8446, f853])).
fof(f8446, plain, ((e3 = op(op(e2, e3), e3)) | ~ spl4_39), inference(forward_demodulation, [], [f348, f773])).
fof(f8552, plain, (~ spl4_8 | ~ spl4_58), inference(avatar_split_clause, [], [f8546, f851, f641])).
fof(f8546, plain, (~ (e2 = op(e4, e3)) | ~ spl4_58), inference(backward_demodulation, [], [f394, f853])).
fof(f8544, plain, (spl4_37 | ~ spl4_69 | ~ spl4_88), inference(avatar_split_clause, [], [f8542, f977, f897, f763])).
fof(f8542, plain, ((e1 = op(e3, e2)) | (~ spl4_69 | ~ spl4_88)), inference(backward_demodulation, [], [f8298, f899])).
fof(f8298, plain, ((e1 = op(op(e2, e1), e2)) | ~ spl4_88), inference(backward_demodulation, [], [f338, f979])).
fof(f8536, plain, (spl4_28 | ~ spl4_75 | ~ spl4_114), inference(avatar_split_clause, [], [f8530, f1086, f922, f725])).
fof(f8530, plain, ((e2 = op(e3, e4)) | (~ spl4_75 | ~ spl4_114)), inference(backward_demodulation, [], [f8491, f924])).
fof(f8491, plain, ((e2 = op(e3, op(e2, e0))) | ~ spl4_114), inference(backward_demodulation, [], [f341, f1088])).
fof(f8535, plain, (spl4_6 | ~ spl4_75 | ~ spl4_114), inference(avatar_split_clause, [], [f8529, f1086, f922, f633])).
fof(f8529, plain, ((e0 = op(e4, e3)) | (~ spl4_75 | ~ spl4_114)), inference(backward_demodulation, [], [f8490, f924])).
fof(f8490, plain, ((e0 = op(op(e2, e0), e3)) | ~ spl4_114), inference(backward_demodulation, [], [f333, f1088])).
fof(f8532, plain, (~ spl4_70 | ~ spl4_75), inference(avatar_split_clause, [], [f8520, f922, f901])).
fof(f8520, plain, (~ (e4 = op(e2, e1)) | ~ spl4_75), inference(backward_demodulation, [], [f426, f924])).
fof(f8531, plain, (~ spl4_25 | ~ spl4_75), inference(avatar_split_clause, [], [f8519, f922, f712])).
fof(f8519, plain, (~ (e4 = op(e4, e0)) | ~ spl4_75), inference(backward_demodulation, [], [f364, f924])).
fof(f8475, plain, (~ spl4_8 | ~ spl4_4 | spl4_139), inference(avatar_split_clause, [], [f8474, f1267, f624, f641])).
fof(f8474, plain, (~ (e2 = op(e4, e3)) | (~ spl4_4 | spl4_139)), inference(forward_demodulation, [], [f1269, f626])).
fof(f8473, plain, (~ spl4_74 | ~ spl4_61 | spl4_140), inference(avatar_split_clause, [], [f8472, f1272, f864, f918])).
fof(f8472, plain, (~ (e3 = op(e2, e0)) | (~ spl4_61 | spl4_140)), inference(forward_demodulation, [], [f1274, f866])).
fof(f8464, plain, (~ spl4_114 | ~ spl4_123 | spl4_168), inference(avatar_split_clause, [], [f8463, f1414, f1124, f1086])).
fof(f8463, plain, (~ (e3 = op(e0, e2)) | (~ spl4_123 | spl4_168)), inference(forward_demodulation, [], [f1416, f1126])).
fof(f8437, plain, (~ spl4_14 | ~ spl4_39), inference(avatar_split_clause, [], [f8436, f771, f666])).
fof(f8436, plain, (~ (e3 = op(e4, e2)) | ~ spl4_39), inference(forward_demodulation, [], [f385, f773])).
fof(f8434, plain, (~ spl4_13 | ~ spl4_88), inference(avatar_split_clause, [], [f8433, f977, f662])).
fof(f8433, plain, (~ (e2 = op(e4, e2)) | ~ spl4_88), inference(forward_demodulation, [], [f382, f979])).
fof(f8432, plain, (~ spl4_4 | ~ spl4_8 | ~ spl4_61 | spl4_133), inference(avatar_contradiction_clause, [], [f8431])).
fof(f8431, plain, ($false | (~ spl4_4 | ~ spl4_8 | ~ spl4_61 | spl4_133)), inference(subsumption_resolution, [], [f8427, f866])).
fof(f8427, plain, (~ (e0 = op(e2, e2)) | (~ spl4_4 | ~ spl4_8 | spl4_133)), inference(backward_demodulation, [], [f8205, f643])).
fof(f8205, plain, (~ (e0 = op(op(e4, e3), op(e4, e3))) | (~ spl4_4 | spl4_133)), inference(forward_demodulation, [], [f1243, f626])).
fof(f1243, plain, (~ (e0 = op(op(e4, op(e4, e4)), op(e4, op(e4, e4)))) | spl4_133), inference(avatar_component_clause, [], [f1241])).
fof(f8413, plain, (spl4_16 | ~ spl4_25 | ~ spl4_102), inference(avatar_split_clause, [], [f8410, f1036, f712, f675])).
fof(f8410, plain, ((e0 = op(e4, e1)) | (~ spl4_25 | ~ spl4_102)), inference(backward_demodulation, [], [f8243, f714])).
fof(f8243, plain, ((e0 = op(op(e4, e0), e1)) | ~ spl4_102), inference(backward_demodulation, [], [f335, f1038])).
fof(f335, plain, (e0 = op(op(e4, e0), op(e0, e4))), inference(cnf_transformation, [], [f3])).
fof(f8412, plain, (~ spl4_20 | ~ spl4_25), inference(avatar_split_clause, [], [f8408, f712, f691])).
fof(f8408, plain, (~ (e4 = op(e4, e1)) | ~ spl4_25), inference(backward_demodulation, [], [f446, f714])).
fof(f8394, plain, (~ spl4_27 | ~ spl4_47), inference(avatar_split_clause, [], [f8388, f805, f721])).
fof(f8388, plain, (~ (e1 = op(e3, e4)) | ~ spl4_47), inference(backward_demodulation, [], [f439, f807])).
fof(f8361, plain, (~ spl4_36 | ~ spl4_61), inference(avatar_split_clause, [], [f8345, f864, f759])).
fof(f759, plain, (spl4_36 <=> (e0 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl4_36])])).
fof(f8345, plain, (~ (e0 = op(e3, e2)) | ~ spl4_61), inference(backward_demodulation, [], [f383, f866])).
fof(f8341, plain, (~ spl4_20 | ~ spl4_70), inference(avatar_split_clause, [], [f8335, f901, f691])).
fof(f8335, plain, (~ (e4 = op(e4, e1)) | ~ spl4_70), inference(backward_demodulation, [], [f374, f903])).
fof(f8334, plain, (spl4_26 | ~ spl4_74 | ~ spl4_115), inference(avatar_split_clause, [], [f8330, f1090, f918, f717])).
fof(f8330, plain, ((e0 = op(e3, e4)) | (~ spl4_74 | ~ spl4_115)), inference(backward_demodulation, [], [f7973, f920])).
fof(f7973, plain, ((e0 = op(op(e2, e0), e4)) | ~ spl4_115), inference(forward_demodulation, [], [f333, f1092])).
fof(f8333, plain, (spl4_8 | ~ spl4_74 | ~ spl4_115), inference(avatar_split_clause, [], [f8329, f1090, f918, f641])).
fof(f8329, plain, ((e2 = op(e4, e3)) | (~ spl4_74 | ~ spl4_115)), inference(backward_demodulation, [], [f7972, f920])).
fof(f7972, plain, ((e2 = op(e4, op(e2, e0))) | ~ spl4_115), inference(forward_demodulation, [], [f341, f1092])).
fof(f8331, plain, (~ spl4_49 | ~ spl4_74), inference(avatar_split_clause, [], [f8322, f918, f813])).
fof(f8322, plain, (~ (e3 = op(e3, e0)) | ~ spl4_74), inference(backward_demodulation, [], [f363, f920])).
fof(f363, plain, ~ (op(e2, e0) = op(e3, e0)), inference(cnf_transformation, [], [f4])).
fof(f8321, plain, (~ spl4_55 | ~ spl4_80), inference(avatar_split_clause, [], [f8319, f943, f838])).
fof(f8319, plain, (~ (e4 = op(e2, e4)) | ~ spl4_80), inference(backward_demodulation, [], [f400, f945])).
fof(f8316, plain, (spl4_57 | ~ spl4_43 | ~ spl4_84), inference(avatar_split_clause, [], [f8313, f960, f788, f847])).
fof(f8313, plain, ((e1 = op(e2, e3)) | (~ spl4_43 | ~ spl4_84)), inference(backward_demodulation, [], [f8186, f962])).
fof(f8186, plain, ((e1 = op(e2, op(e1, e3))) | ~ spl4_43), inference(forward_demodulation, [], [f339, f790])).
fof(f8307, plain, (~ spl4_78 | ~ spl4_88), inference(avatar_split_clause, [], [f8303, f977, f935])).
fof(f8303, plain, (~ (e2 = op(e1, e4)) | ~ spl4_88), inference(backward_demodulation, [], [f424, f979])).
fof(f8306, plain, (~ spl4_63 | ~ spl4_88), inference(avatar_split_clause, [], [f8300, f977, f872])).
fof(f8300, plain, (~ (e2 = op(e2, e2)) | ~ spl4_88), inference(backward_demodulation, [], [f380, f979])).
fof(f8285, plain, (~ spl4_82 | ~ spl4_92), inference(avatar_split_clause, [], [f8272, f994, f952])).
fof(f8272, plain, (~ (e1 = op(e1, e3)) | ~ spl4_92), inference(backward_demodulation, [], [f421, f996])).
fof(f8284, plain, (~ spl4_67 | ~ spl4_92), inference(avatar_split_clause, [], [f8269, f994, f889])).
fof(f8269, plain, (~ (e1 = op(e2, e1)) | ~ spl4_92), inference(backward_demodulation, [], [f370, f996])).
fof(f8267, plain, (spl4_47 | ~ spl4_96 | ~ spl4_119), inference(avatar_split_clause, [], [f8264, f1107, f1011, f805])).
fof(f8264, plain, ((e1 = op(e3, e0)) | (~ spl4_96 | ~ spl4_119)), inference(backward_demodulation, [], [f8189, f1013])).
fof(f8189, plain, ((e1 = op(e3, op(e1, e0))) | ~ spl4_119), inference(forward_demodulation, [], [f336, f1109])).
fof(f8266, plain, (~ spl4_91 | ~ spl4_96), inference(avatar_split_clause, [], [f8255, f1011, f990])).
fof(f8255, plain, (~ (e0 = op(e1, e1)) | ~ spl4_96), inference(backward_demodulation, [], [f416, f1013])).
fof(f8241, plain, (~ spl4_6 | ~ spl4_106), inference(avatar_split_clause, [], [f8235, f1053, f633])).
fof(f8235, plain, (~ (e0 = op(e4, e3)) | ~ spl4_106), inference(backward_demodulation, [], [f389, f1055])).
fof(f8198, plain, (~ spl4_27 | ~ spl4_35 | spl4_149), inference(avatar_split_clause, [], [f8197, f1315, f754, f721])).
fof(f1315, plain, (spl4_149 <=> (e1 = op(e3, op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl4_149])])).
fof(f8197, plain, (~ (e1 = op(e3, e4)) | (~ spl4_35 | spl4_149)), inference(forward_demodulation, [], [f1317, f756])).
fof(f1317, plain, (~ (e1 = op(e3, op(e3, e3))) | spl4_149), inference(avatar_component_clause, [], [f1315])).
fof(f8195, plain, (~ spl4_6 | ~ spl4_4 | spl4_167), inference(avatar_split_clause, [], [f8194, f1409, f624, f633])).
fof(f8194, plain, (~ (e0 = op(e4, e3)) | (~ spl4_4 | spl4_167)), inference(forward_demodulation, [], [f1411, f626])).
fof(f8175, plain, (~ spl4_26 | ~ spl4_35 | spl4_169), inference(avatar_split_clause, [], [f7906, f1419, f754, f717])).
fof(f7906, plain, (~ (e0 = op(e3, e4)) | (~ spl4_35 | spl4_169)), inference(backward_demodulation, [], [f1421, f756])).
fof(f8160, plain, (~ spl4_38 | ~ spl4_43), inference(avatar_split_clause, [], [f8155, f788, f767])).
fof(f8155, plain, (~ (e2 = op(e3, e2)) | ~ spl4_43), inference(backward_demodulation, [], [f440, f790])).
fof(f8119, plain, (spl4_43 | ~ spl4_67 | ~ spl4_89), inference(avatar_split_clause, [], [f8118, f981, f889, f788])).
fof(f8118, plain, ((e2 = op(e3, e1)) | (~ spl4_67 | ~ spl4_89)), inference(backward_demodulation, [], [f8085, f891])).
fof(f8085, plain, ((e2 = op(e3, op(e2, e1))) | ~ spl4_89), inference(backward_demodulation, [], [f342, f983])).
fof(f8083, plain, (~ spl4_4 | ~ spl4_91 | ~ spl4_100 | spl4_184), inference(avatar_contradiction_clause, [], [f8082])).
fof(f8082, plain, ($false | (~ spl4_4 | ~ spl4_91 | ~ spl4_100 | spl4_184)), inference(subsumption_resolution, [], [f8081, f626])).
fof(f8081, plain, (~ (e3 = op(e4, e4)) | (~ spl4_91 | ~ spl4_100 | spl4_184)), inference(forward_demodulation, [], [f8071, f1029])).
fof(f8071, plain, (~ (e3 = op(op(e1, e0), op(e1, e0))) | (~ spl4_91 | spl4_184)), inference(backward_demodulation, [], [f1520, f992])).
fof(f8077, plain, (~ spl4_91 | ~ spl4_100 | spl4_146), inference(avatar_contradiction_clause, [], [f8076])).
fof(f8076, plain, ($false | (~ spl4_91 | ~ spl4_100 | spl4_146)), inference(subsumption_resolution, [], [f8065, f1029])).
fof(f8065, plain, (~ (e4 = op(e1, e0)) | (~ spl4_91 | spl4_146)), inference(backward_demodulation, [], [f1302, f992])).
fof(f8051, plain, (~ spl4_76 | ~ spl4_101), inference(avatar_split_clause, [], [f8044, f1032, f927])).
fof(f8044, plain, (~ (e0 = op(e1, e4)) | ~ spl4_101), inference(backward_demodulation, [], [f396, f1034])).
fof(f8032, plain, (~ spl4_44 | ~ spl4_119), inference(avatar_split_clause, [], [f8021, f1107, f792])).
fof(f8021, plain, (~ (e3 = op(e3, e1)) | ~ spl4_119), inference(backward_demodulation, [], [f368, f1109])).
fof(f8017, plain, (~ spl4_72 | ~ spl4_122), inference(avatar_split_clause, [], [f8002, f1120, f910])).
fof(f8002, plain, (~ (e1 = op(e2, e0)) | ~ spl4_122), inference(backward_demodulation, [], [f357, f1122])).
fof(f7984, plain, (~ spl4_123 | ~ spl4_4 | ~ spl4_6 | spl4_175), inference(avatar_split_clause, [], [f7983, f1457, f633, f624, f1124])).
fof(f7983, plain, (~ (op(e0, e0) = e2) | (~ spl4_4 | ~ spl4_6 | spl4_175)), inference(forward_demodulation, [], [f7982, f635])).
fof(f7982, plain, (~ (e2 = op(op(e4, e3), op(e4, e3))) | (~ spl4_4 | spl4_175)), inference(forward_demodulation, [], [f1459, f626])).
fof(f7955, plain, (spl4_100 | ~ spl4_6 | ~ spl4_27), inference(avatar_split_clause, [], [f7931, f721, f633, f1027])).
fof(f7931, plain, ((e4 = op(e1, e0)) | (~ spl4_6 | ~ spl4_27)), inference(backward_demodulation, [], [f7910, f635])).
fof(f7910, plain, ((e4 = op(e1, op(e4, e3))) | ~ spl4_27), inference(backward_demodulation, [], [f354, f723])).
fof(f7932, plain, (~ spl4_1 | ~ spl4_6), inference(avatar_split_clause, [], [f7925, f633, f612])).
fof(f7925, plain, (~ (e0 = op(e4, e4)) | ~ spl4_6), inference(backward_demodulation, [], [f455, f635])).
fof(f7901, plain, (~ spl4_38 | ~ spl4_59), inference(avatar_contradiction_clause, [], [f7900])).
fof(f7900, plain, ($false | (~ spl4_38 | ~ spl4_59)), inference(subsumption_resolution, [], [f7899, f463])).
fof(f7899, plain, ((e2 = e3) | (~ spl4_38 | ~ spl4_59)), inference(forward_demodulation, [], [f7896, f769])).
fof(f7896, plain, ((e3 = op(e3, e2)) | (~ spl4_38 | ~ spl4_59)), inference(backward_demodulation, [], [f7569, f769])).
fof(f7569, plain, ((e3 = op(e3, op(e3, e2))) | ~ spl4_59), inference(backward_demodulation, [], [f348, f857])).
fof(f7846, plain, (spl4_63 | ~ spl4_68 | ~ spl4_88), inference(avatar_split_clause, [], [f7842, f977, f893, f872])).
fof(f7842, plain, ((e2 = op(e2, e2)) | (~ spl4_68 | ~ spl4_88)), inference(backward_demodulation, [], [f7789, f895])).
fof(f7789, plain, ((e2 = op(e2, op(e2, e1))) | ~ spl4_88), inference(backward_demodulation, [], [f342, f979])).
fof(f7836, plain, (spl4_43 | ~ spl4_72 | ~ spl4_114), inference(avatar_contradiction_clause, [], [f7835])).
fof(f7835, plain, ($false | (spl4_43 | ~ spl4_72 | ~ spl4_114)), inference(subsumption_resolution, [], [f7831, f789])).
fof(f789, plain, (~ (e2 = op(e3, e1)) | spl4_43), inference(avatar_component_clause, [], [f788])).
fof(f7831, plain, ((e2 = op(e3, e1)) | (~ spl4_72 | ~ spl4_114)), inference(backward_demodulation, [], [f7452, f912])).
fof(f7452, plain, ((e2 = op(e3, op(e2, e0))) | ~ spl4_114), inference(backward_demodulation, [], [f341, f1088])).
fof(f7834, plain, (~ spl4_72 | spl4_81 | ~ spl4_114), inference(avatar_contradiction_clause, [], [f7833])).
fof(f7833, plain, ($false | (~ spl4_72 | spl4_81 | ~ spl4_114)), inference(subsumption_resolution, [], [f7830, f949])).
fof(f949, plain, (~ (e0 = op(e1, e3)) | spl4_81), inference(avatar_component_clause, [], [f948])).
fof(f7830, plain, ((e0 = op(e1, e3)) | (~ spl4_72 | ~ spl4_114)), inference(backward_demodulation, [], [f7451, f912])).
fof(f7451, plain, ((e0 = op(op(e2, e0), e3)) | ~ spl4_114), inference(backward_demodulation, [], [f333, f1088])).
fof(f7822, plain, (~ spl4_20 | spl4_22 | ~ spl4_76), inference(avatar_contradiction_clause, [], [f7821])).
fof(f7821, plain, ($false | (~ spl4_20 | spl4_22 | ~ spl4_76)), inference(subsumption_resolution, [], [f7816, f701])).
fof(f701, plain, (~ (e1 = op(e4, e0)) | spl4_22), inference(avatar_component_clause, [], [f700])).
fof(f7816, plain, ((e1 = op(e4, e0)) | (~ spl4_20 | ~ spl4_76)), inference(backward_demodulation, [], [f7653, f929])).
fof(f7653, plain, ((e1 = op(e4, op(e1, e4))) | ~ spl4_20), inference(forward_demodulation, [], [f340, f693])).
fof(f7798, plain, (~ spl4_83 | ~ spl4_88), inference(avatar_split_clause, [], [f7792, f977, f956])).
fof(f7792, plain, (~ (e2 = op(e1, e3)) | ~ spl4_88), inference(backward_demodulation, [], [f423, f979])).
fof(f7754, plain, (~ spl4_99 | ~ spl4_100), inference(avatar_contradiction_clause, [], [f7753])).
fof(f7753, plain, ($false | (~ spl4_99 | ~ spl4_100)), inference(subsumption_resolution, [], [f7752, f465])).
fof(f7752, plain, ((e3 = e4) | (~ spl4_99 | ~ spl4_100)), inference(backward_demodulation, [], [f1029, f1025])).
fof(f7751, plain, (~ spl4_100 | spl4_102 | ~ spl4_116), inference(avatar_contradiction_clause, [], [f7750])).
fof(f7750, plain, ($false | (~ spl4_100 | spl4_102 | ~ spl4_116)), inference(subsumption_resolution, [], [f7746, f1037])).
fof(f1037, plain, (~ (e1 = op(e0, e4)) | spl4_102), inference(avatar_component_clause, [], [f1036])).
fof(f7746, plain, ((e1 = op(e0, e4)) | (~ spl4_100 | ~ spl4_116)), inference(backward_demodulation, [], [f7704, f1029])).
fof(f7704, plain, ((e1 = op(e0, op(e1, e0))) | ~ spl4_116), inference(backward_demodulation, [], [f336, f1097])).
fof(f7747, plain, (~ spl4_85 | ~ spl4_100), inference(avatar_split_clause, [], [f7743, f1027, f964])).
fof(f7743, plain, (~ (e4 = op(e1, e3)) | ~ spl4_100), inference(backward_demodulation, [], [f418, f1029])).
fof(f7737, plain, (~ spl4_78 | ~ spl4_103), inference(avatar_split_clause, [], [f7732, f1040, f935])).
fof(f7732, plain, (~ (e2 = op(e1, e4)) | ~ spl4_103), inference(backward_demodulation, [], [f396, f1042])).
fof(f7718, plain, (~ spl4_114 | ~ spl4_115), inference(avatar_contradiction_clause, [], [f7717])).
fof(f7717, plain, ($false | (~ spl4_114 | ~ spl4_115)), inference(subsumption_resolution, [], [f7716, f465])).
fof(f7716, plain, ((e3 = e4) | (~ spl4_114 | ~ spl4_115)), inference(forward_demodulation, [], [f1092, f1088])).
fof(f7714, plain, (~ spl4_41 | ~ spl4_116), inference(avatar_split_clause, [], [f7707, f1095, f780])).
fof(f7707, plain, (~ (e0 = op(e3, e1)) | ~ spl4_116), inference(backward_demodulation, [], [f368, f1097])).
fof(f7697, plain, (~ spl4_32 | ~ spl4_114 | ~ spl4_123 | spl4_164), inference(avatar_split_clause, [], [f7696, f1396, f1124, f1086, f742])).
fof(f1396, plain, (spl4_164 <=> (e1 = op(op(e0, op(e0, e0)), op(e0, op(e0, e0))))), introduced(avatar_definition, [new_symbols(naming, [spl4_164])])).
fof(f7696, plain, (~ (e1 = op(e3, e3)) | (~ spl4_114 | ~ spl4_123 | spl4_164)), inference(forward_demodulation, [], [f7685, f1088])).
fof(f7685, plain, (~ (e1 = op(op(e0, e2), op(e0, e2))) | (~ spl4_123 | spl4_164)), inference(backward_demodulation, [], [f1398, f1126])).
fof(f1398, plain, (~ (e1 = op(op(e0, op(e0, e0)), op(e0, op(e0, e0)))) | spl4_164), inference(avatar_component_clause, [], [f1396])).
fof(f7695, plain, (~ spl4_103 | ~ spl4_123), inference(avatar_split_clause, [], [f7682, f1124, f1040])).
fof(f7682, plain, (~ (e2 = op(e0, e4)) | ~ spl4_123), inference(backward_demodulation, [], [f409, f1126])).
fof(f7692, plain, (~ spl4_48 | ~ spl4_123), inference(avatar_split_clause, [], [f7678, f1124, f809])).
fof(f7678, plain, (~ (e2 = op(e3, e0)) | ~ spl4_123), inference(backward_demodulation, [], [f358, f1126])).
fof(f7671, plain, (~ spl4_41 | ~ spl4_149 | spl4_159), inference(avatar_split_clause, [], [f7668, f1368, f1315, f780])).
fof(f1368, plain, (spl4_159 <=> (e0 = op(e3, op(e3, op(e3, e3))))), introduced(avatar_definition, [new_symbols(naming, [spl4_159])])).
fof(f7668, plain, (~ (e0 = op(e3, e1)) | (~ spl4_149 | spl4_159)), inference(backward_demodulation, [], [f1370, f1316])).
fof(f1316, plain, ((e1 = op(e3, op(e3, e3))) | ~ spl4_149), inference(avatar_component_clause, [], [f1315])).
fof(f1370, plain, (~ (e0 = op(e3, op(e3, op(e3, e3)))) | spl4_159), inference(avatar_component_clause, [], [f1368])).
fof(f7650, plain, (spl4_78 | ~ spl4_12 | ~ spl4_55), inference(avatar_split_clause, [], [f7617, f838, f658, f935])).
fof(f7617, plain, ((e2 = op(e1, e4)) | (~ spl4_12 | ~ spl4_55)), inference(backward_demodulation, [], [f7574, f660])).
fof(f7574, plain, ((e2 = op(op(e4, e2), e4)) | ~ spl4_55), inference(backward_demodulation, [], [f345, f840])).
fof(f7635, plain, (spl4_8 | ~ spl4_135 | ~ spl4_143), inference(avatar_split_clause, [], [f7632, f1287, f1249, f641])).
fof(f7632, plain, ((e2 = op(e4, e3)) | (~ spl4_135 | ~ spl4_143)), inference(backward_demodulation, [], [f1288, f1250])).
fof(f1250, plain, ((e3 = op(e4, op(e4, e4))) | ~ spl4_135), inference(avatar_component_clause, [], [f1249])).
fof(f1288, plain, ((e2 = op(e4, op(e4, op(e4, e4)))) | ~ spl4_143), inference(avatar_component_clause, [], [f1287])).
fof(f7634, plain, (~ spl4_33 | ~ spl4_135 | spl4_175), inference(avatar_split_clause, [], [f7631, f1457, f1249, f746])).
fof(f7631, plain, (~ (e2 = op(e3, e3)) | (~ spl4_135 | spl4_175)), inference(backward_demodulation, [], [f1459, f1250])).
fof(f7629, plain, (~ spl4_2 | ~ spl4_12), inference(avatar_split_clause, [], [f7628, f658, f616])).
fof(f7628, plain, (~ (e1 = op(e4, e4)) | ~ spl4_12), inference(forward_demodulation, [], [f454, f660])).
fof(f7619, plain, (~ spl4_7 | ~ spl4_12), inference(avatar_split_clause, [], [f7613, f658, f637])).
fof(f7613, plain, (~ (e1 = op(e4, e3)) | ~ spl4_12), inference(backward_demodulation, [], [f453, f660])).
fof(f7612, plain, (spl4_12 | ~ spl4_20 | ~ spl4_78), inference(avatar_split_clause, [], [f7609, f935, f691, f658])).
fof(f7609, plain, ((e1 = op(e4, e2)) | (~ spl4_20 | ~ spl4_78)), inference(backward_demodulation, [], [f6770, f693])).
fof(f6770, plain, ((e1 = op(op(e4, e1), e2)) | ~ spl4_78), inference(backward_demodulation, [], [f340, f937])).
fof(f7611, plain, (~ spl4_15 | ~ spl4_20), inference(avatar_split_clause, [], [f7607, f691, f670])).
fof(f7607, plain, (~ (e4 = op(e4, e2)) | ~ spl4_20), inference(backward_demodulation, [], [f450, f693])).
fof(f7606, plain, (~ spl4_7 | ~ spl4_32), inference(avatar_split_clause, [], [f7601, f742, f637])).
fof(f7601, plain, (~ (e1 = op(e4, e3)) | ~ spl4_32), inference(backward_demodulation, [], [f395, f744])).
fof(f7599, plain, (spl4_8 | ~ spl4_40 | ~ spl4_59), inference(avatar_split_clause, [], [f7596, f855, f775, f641])).
fof(f7596, plain, ((e2 = op(e4, e3)) | (~ spl4_40 | ~ spl4_59)), inference(backward_demodulation, [], [f7568, f777])).
fof(f7568, plain, ((e2 = op(op(e3, e2), e3)) | ~ spl4_59), inference(backward_demodulation, [], [f344, f857])).
fof(f7593, plain, (~ spl4_36 | ~ spl4_41), inference(avatar_split_clause, [], [f7585, f780, f759])).
fof(f7585, plain, (~ (e0 = op(e3, e2)) | ~ spl4_41), inference(backward_demodulation, [], [f440, f782])).
fof(f7583, plain, (~ spl4_33 | ~ spl4_48), inference(avatar_split_clause, [], [f7578, f809, f746])).
fof(f7578, plain, (~ (e2 = op(e3, e3)) | ~ spl4_48), inference(backward_demodulation, [], [f438, f811])).
fof(f7545, plain, (~ spl4_52 | ~ spl4_67), inference(avatar_split_clause, [], [f7539, f889, f826])).
fof(f7539, plain, (~ (e1 = op(e2, e4)) | ~ spl4_67), inference(backward_demodulation, [], [f432, f891])).
fof(f7534, plain, (spl4_48 | ~ spl4_71 | ~ spl4_114), inference(avatar_split_clause, [], [f7532, f1086, f906, f809])).
fof(f7532, plain, ((e2 = op(e3, e0)) | (~ spl4_71 | ~ spl4_114)), inference(backward_demodulation, [], [f7452, f908])).
fof(f7514, plain, (~ spl4_36 | ~ spl4_86), inference(avatar_split_clause, [], [f7507, f969, f759])).
fof(f7507, plain, (~ (e0 = op(e3, e2)) | ~ spl4_86), inference(backward_demodulation, [], [f381, f971])).
fof(f7502, plain, (spl4_85 | ~ spl4_94 | ~ spl4_146), inference(avatar_split_clause, [], [f7497, f1300, f1002, f964])).
fof(f7497, plain, ((e4 = op(e1, e3)) | (~ spl4_94 | ~ spl4_146)), inference(backward_demodulation, [], [f1301, f1004])).
fof(f7487, plain, (~ spl4_92 | ~ spl4_97), inference(avatar_split_clause, [], [f7478, f1015, f994])).
fof(f7478, plain, (~ (e1 = op(e1, e1)) | ~ spl4_97), inference(backward_demodulation, [], [f416, f1017])).
fof(f7474, plain, (spl4_85 | ~ spl4_24 | ~ spl4_102), inference(avatar_split_clause, [], [f7467, f1036, f708, f964])).
fof(f7467, plain, ((e4 = op(e1, e3)) | (~ spl4_24 | ~ spl4_102)), inference(backward_demodulation, [], [f7413, f1038])).
fof(f7413, plain, ((e4 = op(op(e0, e4), e3)) | ~ spl4_24), inference(forward_demodulation, [], [f351, f710])).
fof(f7473, plain, (~ spl4_102 | ~ spl4_118 | ~ spl4_125 | spl4_165), inference(avatar_contradiction_clause, [], [f7472])).
fof(f7472, plain, ($false | (~ spl4_102 | ~ spl4_118 | ~ spl4_125 | spl4_165)), inference(subsumption_resolution, [], [f7465, f1105])).
fof(f7465, plain, (~ (e2 = op(e0, e1)) | (~ spl4_102 | ~ spl4_125 | spl4_165)), inference(backward_demodulation, [], [f7198, f1038])).
fof(f7198, plain, (~ (e2 = op(e0, op(e0, e4))) | (~ spl4_125 | spl4_165)), inference(backward_demodulation, [], [f1402, f1134])).
fof(f7470, plain, (~ spl4_52 | ~ spl4_102), inference(avatar_split_clause, [], [f7460, f1036, f826])).
fof(f7460, plain, (~ (e1 = op(e2, e4)) | ~ spl4_102), inference(backward_demodulation, [], [f397, f1038])).
fof(f7438, plain, (~ spl4_48 | spl4_142 | ~ spl4_169), inference(avatar_split_clause, [], [f7437, f1419, f1282, f809])).
fof(f1282, plain, (spl4_142 <=> (e2 = op(e3, op(e3, op(e3, e3))))), introduced(avatar_definition, [new_symbols(naming, [spl4_142])])).
fof(f7437, plain, (~ (e2 = op(e3, e0)) | (spl4_142 | ~ spl4_169)), inference(forward_demodulation, [], [f1284, f1420])).
fof(f1284, plain, (~ (e2 = op(e3, op(e3, op(e3, e3)))) | spl4_142), inference(avatar_component_clause, [], [f1282])).
fof(f7430, plain, (~ spl4_78 | spl4_145 | ~ spl4_146), inference(avatar_contradiction_clause, [], [f7429])).
fof(f7429, plain, ($false | (~ spl4_78 | spl4_145 | ~ spl4_146)), inference(subsumption_resolution, [], [f7425, f937])).
fof(f7425, plain, (~ (e2 = op(e1, e4)) | (spl4_145 | ~ spl4_146)), inference(backward_demodulation, [], [f1298, f1301])).
fof(f7409, plain, (~ spl4_75 | ~ spl4_125), inference(avatar_split_clause, [], [f7408, f1132, f922])).
fof(f7408, plain, (~ (e4 = op(e2, e0)) | ~ spl4_125), inference(forward_demodulation, [], [f357, f1134])).
fof(f7395, plain, (~ spl4_32 | ~ spl4_135 | spl4_160), inference(avatar_split_clause, [], [f7186, f1373, f1249, f742])).
fof(f1373, plain, (spl4_160 <=> (e1 = op(op(e4, op(e4, e4)), op(e4, op(e4, e4))))), introduced(avatar_definition, [new_symbols(naming, [spl4_160])])).
fof(f7186, plain, (~ (e1 = op(e3, e3)) | (~ spl4_135 | spl4_160)), inference(backward_demodulation, [], [f1375, f1250])).
fof(f1375, plain, (~ (e1 = op(op(e4, op(e4, e4)), op(e4, op(e4, e4)))) | spl4_160), inference(avatar_component_clause, [], [f1373])).
fof(f7376, plain, (~ spl4_5 | ~ spl4_15), inference(avatar_split_clause, [], [f7373, f670, f628])).
fof(f7373, plain, (~ (e4 = op(e4, e4)) | ~ spl4_15), inference(backward_demodulation, [], [f454, f672])).
fof(f7368, plain, (~ spl4_18 | spl4_62 | ~ spl4_78), inference(avatar_contradiction_clause, [], [f7367])).
fof(f7367, plain, ($false | (~ spl4_18 | spl4_62 | ~ spl4_78)), inference(subsumption_resolution, [], [f7364, f869])).
fof(f869, plain, (~ (e1 = op(e2, e2)) | spl4_62), inference(avatar_component_clause, [], [f868])).
fof(f7364, plain, ((e1 = op(e2, e2)) | (~ spl4_18 | ~ spl4_78)), inference(backward_demodulation, [], [f6770, f685])).
fof(f7358, plain, (~ spl4_19 | ~ spl4_24), inference(avatar_split_clause, [], [f7351, f708, f687])).
fof(f7351, plain, (~ (e3 = op(e4, e1)) | ~ spl4_24), inference(backward_demodulation, [], [f446, f710])).
fof(f7350, plain, (spl4_7 | ~ spl4_45 | ~ spl4_84), inference(avatar_split_clause, [], [f7348, f960, f796, f637])).
fof(f7348, plain, ((e1 = op(e4, e3)) | (~ spl4_45 | ~ spl4_84)), inference(backward_demodulation, [], [f7307, f798])).
fof(f7307, plain, ((e1 = op(op(e3, e1), e3)) | ~ spl4_84), inference(backward_demodulation, [], [f339, f962])).
fof(f7346, plain, (~ spl4_42 | ~ spl4_47), inference(avatar_split_clause, [], [f7339, f805, f784])).
fof(f7339, plain, (~ (e1 = op(e3, e1)) | ~ spl4_47), inference(backward_demodulation, [], [f436, f807])).
fof(f7336, plain, (spl4_24 | ~ spl4_36 | ~ spl4_60), inference(avatar_split_clause, [], [f7328, f859, f759, f708])).
fof(f7328, plain, ((e3 = op(e4, e0)) | (~ spl4_36 | ~ spl4_60)), inference(backward_demodulation, [], [f7151, f861])).
fof(f7151, plain, ((e3 = op(op(e2, e3), e0)) | ~ spl4_36), inference(forward_demodulation, [], [f348, f761])).
fof(f761, plain, ((e0 = op(e3, e2)) | ~ spl4_36), inference(avatar_component_clause, [], [f759])).
fof(f7314, plain, (~ spl4_78 | ~ spl4_80), inference(avatar_contradiction_clause, [], [f7313])).
fof(f7313, plain, ($false | (~ spl4_78 | ~ spl4_80)), inference(subsumption_resolution, [], [f7312, f464])).
fof(f464, plain, ~ (e2 = e4), inference(cnf_transformation, [], [f5])).
fof(f7312, plain, ((e2 = e4) | (~ spl4_78 | ~ spl4_80)), inference(forward_demodulation, [], [f945, f937])).
fof(f7283, plain, (~ spl4_87 | ~ spl4_92), inference(avatar_split_clause, [], [f7269, f994, f973])).
fof(f7269, plain, (~ (e1 = op(e1, e2)) | ~ spl4_92), inference(backward_demodulation, [], [f420, f996])).
fof(f7228, plain, (spl4_88 | ~ spl4_73 | ~ spl4_112), inference(avatar_split_clause, [], [f7225, f1078, f914, f977])).
fof(f7225, plain, ((e2 = op(e1, e2)) | (~ spl4_73 | ~ spl4_112)), inference(backward_demodulation, [], [f7163, f1080])).
fof(f7163, plain, ((e2 = op(op(e0, e2), e2)) | ~ spl4_73), inference(forward_demodulation, [], [f341, f916])).
fof(f7217, plain, (~ spl4_109 | ~ spl4_119), inference(avatar_split_clause, [], [f7213, f1107, f1065])).
fof(f7213, plain, (~ (e3 = op(e0, e3)) | ~ spl4_119), inference(backward_demodulation, [], [f411, f1109])).
fof(f411, plain, ~ (op(e0, e1) = op(e0, e3)), inference(cnf_transformation, [], [f4])).
fof(f7216, plain, (~ spl4_19 | ~ spl4_119), inference(avatar_split_clause, [], [f7211, f1107, f687])).
fof(f7211, plain, (~ (e3 = op(e4, e1)) | ~ spl4_119), inference(backward_demodulation, [], [f369, f1109])).
fof(f7203, plain, (~ spl4_50 | ~ spl4_125), inference(avatar_split_clause, [], [f7191, f1132, f817])).
fof(f7191, plain, (~ (e4 = op(e3, e0)) | ~ spl4_125), inference(backward_demodulation, [], [f358, f1134])).
fof(f7188, plain, (~ spl4_7 | spl4_134 | ~ spl4_135), inference(avatar_split_clause, [], [f7187, f1249, f1245, f637])).
fof(f7187, plain, (~ (e1 = op(e4, e3)) | (spl4_134 | ~ spl4_135)), inference(forward_demodulation, [], [f1247, f1250])).
fof(f7178, plain, (~ spl4_57 | ~ spl4_64 | spl4_154), inference(avatar_split_clause, [], [f7177, f1341, f876, f847])).
fof(f7177, plain, (~ (e1 = op(e2, e3)) | (~ spl4_64 | spl4_154)), inference(forward_demodulation, [], [f1343, f878])).
fof(f1343, plain, (~ (e1 = op(e2, op(e2, e2))) | spl4_154), inference(avatar_component_clause, [], [f1341])).
fof(f7129, plain, (~ spl4_13 | ~ spl4_52 | spl4_68), inference(avatar_contradiction_clause, [], [f7128])).
fof(f7128, plain, ($false | (~ spl4_13 | ~ spl4_52 | spl4_68)), inference(subsumption_resolution, [], [f7126, f894])).
fof(f894, plain, (~ (e2 = op(e2, e1)) | spl4_68), inference(avatar_component_clause, [], [f893])).
fof(f7126, plain, ((e2 = op(e2, e1)) | (~ spl4_13 | ~ spl4_52)), inference(backward_demodulation, [], [f7079, f664])).
fof(f7079, plain, ((e2 = op(op(e4, e2), e1)) | ~ spl4_52), inference(backward_demodulation, [], [f345, f828])).
fof(f7119, plain, (~ spl4_22 | spl4_96 | ~ spl4_101), inference(avatar_contradiction_clause, [], [f7118])).
fof(f7118, plain, ($false | (~ spl4_22 | spl4_96 | ~ spl4_101)), inference(subsumption_resolution, [], [f7115, f1012])).
fof(f1012, plain, (~ (e0 = op(e1, e0)) | spl4_96), inference(avatar_component_clause, [], [f1011])).
fof(f7115, plain, ((e0 = op(e1, e0)) | (~ spl4_22 | ~ spl4_101)), inference(backward_demodulation, [], [f6983, f702])).
fof(f6983, plain, ((e0 = op(op(e4, e0), e0)) | ~ spl4_101), inference(backward_demodulation, [], [f335, f1034])).
fof(f7117, plain, (~ spl4_2 | ~ spl4_22), inference(avatar_split_clause, [], [f7111, f700, f616])).
fof(f7111, plain, (~ (e1 = op(e4, e4)) | ~ spl4_22), inference(backward_demodulation, [], [f449, f702])).
fof(f7094, plain, (~ spl4_42 | spl4_77 | ~ spl4_85), inference(avatar_contradiction_clause, [], [f7093])).
fof(f7093, plain, ($false | (~ spl4_42 | spl4_77 | ~ spl4_85)), inference(subsumption_resolution, [], [f7090, f932])).
fof(f932, plain, (~ (e1 = op(e1, e4)) | spl4_77), inference(avatar_component_clause, [], [f931])).
fof(f7090, plain, ((e1 = op(e1, e4)) | (~ spl4_42 | ~ spl4_85)), inference(backward_demodulation, [], [f7032, f786])).
fof(f7032, plain, ((e1 = op(op(e3, e1), e4)) | ~ spl4_85), inference(backward_demodulation, [], [f339, f966])).
fof(f7085, plain, (~ spl4_2 | ~ spl4_52), inference(avatar_split_clause, [], [f7081, f826, f616])).
fof(f7081, plain, (~ (e1 = op(e4, e4)) | ~ spl4_52), inference(backward_demodulation, [], [f404, f828])).
fof(f7073, plain, (~ spl4_57 | ~ spl4_64 | ~ spl4_91 | spl4_136), inference(avatar_contradiction_clause, [], [f7072])).
fof(f7072, plain, ($false | (~ spl4_57 | ~ spl4_64 | ~ spl4_91 | spl4_136)), inference(subsumption_resolution, [], [f7066, f992])).
fof(f7066, plain, (~ (e0 = op(e1, e1)) | (~ spl4_57 | ~ spl4_64 | spl4_136)), inference(backward_demodulation, [], [f6537, f849])).
fof(f6537, plain, (~ (e0 = op(op(e2, e3), op(e2, e3))) | (~ spl4_64 | spl4_136)), inference(forward_demodulation, [], [f1256, f878])).
fof(f1256, plain, (~ (e0 = op(op(e2, op(e2, e2)), op(e2, op(e2, e2)))) | spl4_136), inference(avatar_component_clause, [], [f1254])).
fof(f1254, plain, (spl4_136 <=> (e0 = op(op(e2, op(e2, e2)), op(e2, op(e2, e2))))), introduced(avatar_definition, [new_symbols(naming, [spl4_136])])).
fof(f7058, plain, (~ spl4_60 | ~ spl4_70), inference(avatar_split_clause, [], [f7054, f901, f859])).
fof(f7054, plain, (~ (e4 = op(e2, e3)) | ~ spl4_70), inference(backward_demodulation, [], [f431, f903])).
fof(f7017, plain, (~ spl4_85 | ~ spl4_91 | ~ spl4_99 | spl4_157), inference(avatar_split_clause, [], [f7016, f1356, f1023, f990, f964])).
fof(f7016, plain, (~ (e4 = op(e1, e3)) | (~ spl4_91 | ~ spl4_99 | spl4_157)), inference(forward_demodulation, [], [f7006, f1025])).
fof(f7006, plain, (~ (e4 = op(e1, op(e1, e0))) | (~ spl4_91 | spl4_157)), inference(backward_demodulation, [], [f1358, f992])).
fof(f7014, plain, (~ spl4_91 | ~ spl4_99 | spl4_148), inference(avatar_contradiction_clause, [], [f7013])).
fof(f7013, plain, ($false | (~ spl4_91 | ~ spl4_99 | spl4_148)), inference(subsumption_resolution, [], [f7004, f1025])).
fof(f7004, plain, (~ (e3 = op(e1, e0)) | (~ spl4_91 | spl4_148)), inference(backward_demodulation, [], [f1312, f992])).
fof(f7008, plain, (~ spl4_41 | ~ spl4_91), inference(avatar_split_clause, [], [f6994, f990, f780])).
fof(f6994, plain, (~ (e0 = op(e3, e1)) | ~ spl4_91), inference(backward_demodulation, [], [f371, f992])).
fof(f6991, plain, (~ spl4_51 | ~ spl4_101), inference(avatar_split_clause, [], [f6985, f1032, f822])).
fof(f6985, plain, (~ (e0 = op(e2, e4)) | ~ spl4_101), inference(backward_demodulation, [], [f397, f1034])).
fof(f6970, plain, (spl4_57 | ~ spl4_99 | ~ spl4_118), inference(avatar_split_clause, [], [f6962, f1103, f1023, f847])).
fof(f6962, plain, ((e1 = op(e2, e3)) | (~ spl4_99 | ~ spl4_118)), inference(backward_demodulation, [], [f6927, f1105])).
fof(f6927, plain, ((e1 = op(op(e0, e1), e3)) | ~ spl4_99), inference(forward_demodulation, [], [f336, f1025])).
fof(f6954, plain, (~ spl4_117 | ~ spl4_122), inference(avatar_split_clause, [], [f6941, f1120, f1099])).
fof(f6941, plain, (~ (e1 = op(e0, e1)) | ~ spl4_122), inference(backward_demodulation, [], [f406, f1122])).
fof(f6953, plain, (~ spl4_22 | ~ spl4_122), inference(avatar_split_clause, [], [f6940, f1120, f700])).
fof(f6940, plain, (~ (e1 = op(e4, e0)) | ~ spl4_122), inference(backward_demodulation, [], [f359, f1122])).
fof(f6913, plain, (spl4_60 | ~ spl4_19 | ~ spl4_78), inference(avatar_split_clause, [], [f6885, f935, f687, f859])).
fof(f6885, plain, ((e4 = op(e2, e3)) | (~ spl4_19 | ~ spl4_78)), inference(backward_demodulation, [], [f6771, f689])).
fof(f6771, plain, ((e4 = op(e2, op(e4, e1))) | ~ spl4_78), inference(backward_demodulation, [], [f352, f937])).
fof(f6912, plain, (~ spl4_58 | ~ spl4_33), inference(avatar_split_clause, [], [f6484, f746, f851])).
fof(f6484, plain, (~ (e2 = op(e2, e3)) | ~ spl4_33), inference(forward_demodulation, [], [f393, f748])).
fof(f6907, plain, (spl4_37 | ~ spl4_19 | ~ spl4_78), inference(avatar_split_clause, [], [f6884, f935, f687, f763])).
fof(f6884, plain, ((e1 = op(e3, e2)) | (~ spl4_19 | ~ spl4_78)), inference(backward_demodulation, [], [f6770, f689])).
fof(f6902, plain, (~ spl4_4 | ~ spl4_19), inference(avatar_split_clause, [], [f6901, f687, f624])).
fof(f6901, plain, (~ (e3 = op(e4, e4)) | ~ spl4_19), inference(forward_demodulation, [], [f452, f689])).
fof(f6900, plain, (~ spl4_4 | ~ spl4_29), inference(avatar_split_clause, [], [f6899, f729, f624])).
fof(f729, plain, (spl4_29 <=> (e3 = op(e3, e4))), introduced(avatar_definition, [new_symbols(naming, [spl4_29])])).
fof(f6899, plain, (~ (e3 = op(e4, e4)) | ~ spl4_29), inference(forward_demodulation, [], [f405, f731])).
fof(f731, plain, ((e3 = op(e3, e4)) | ~ spl4_29), inference(avatar_component_clause, [], [f729])).
fof(f6875, plain, (spl4_15 | ~ spl4_23 | ~ spl4_105), inference(avatar_split_clause, [], [f6873, f1048, f704, f670])).
fof(f6873, plain, ((e4 = op(e4, e2)) | (~ spl4_23 | ~ spl4_105)), inference(backward_demodulation, [], [f6693, f706])).
fof(f6693, plain, ((e4 = op(e4, op(e4, e0))) | ~ spl4_105), inference(forward_demodulation, [], [f351, f1050])).
fof(f6863, plain, (spl4_19 | ~ spl4_37 | ~ spl4_60), inference(avatar_split_clause, [], [f6862, f859, f763, f687])).
fof(f6862, plain, ((e3 = op(e4, e1)) | (~ spl4_37 | ~ spl4_60)), inference(backward_demodulation, [], [f6812, f765])).
fof(f6812, plain, ((e3 = op(e4, op(e3, e2))) | ~ spl4_60), inference(backward_demodulation, [], [f348, f861])).
fof(f6842, plain, (spl4_29 | ~ spl4_50 | ~ spl4_109), inference(avatar_split_clause, [], [f6839, f1065, f817, f729])).
fof(f6839, plain, ((e3 = op(e3, e4)) | (~ spl4_50 | ~ spl4_109)), inference(backward_demodulation, [], [f5995, f819])).
fof(f5995, plain, ((e3 = op(e3, op(e3, e0))) | ~ spl4_109), inference(backward_demodulation, [], [f346, f1067])).
fof(f6841, plain, (spl4_6 | ~ spl4_50 | ~ spl4_109), inference(avatar_split_clause, [], [f6838, f1065, f817, f633])).
fof(f6838, plain, ((e0 = op(e4, e3)) | (~ spl4_50 | ~ spl4_109)), inference(backward_demodulation, [], [f5994, f819])).
fof(f5994, plain, ((e0 = op(op(e3, e0), e3)) | ~ spl4_109), inference(backward_demodulation, [], [f334, f1067])).
fof(f6800, plain, (spl4_68 | ~ spl4_72 | ~ spl4_113), inference(avatar_split_clause, [], [f6797, f1082, f910, f893])).
fof(f6797, plain, ((e2 = op(e2, e1)) | (~ spl4_72 | ~ spl4_113)), inference(backward_demodulation, [], [f6565, f912])).
fof(f6565, plain, ((e2 = op(e2, op(e2, e0))) | ~ spl4_113), inference(backward_demodulation, [], [f341, f1084])).
fof(f6798, plain, (~ spl4_47 | ~ spl4_72), inference(avatar_split_clause, [], [f6788, f910, f805])).
fof(f6788, plain, (~ (e1 = op(e3, e0)) | ~ spl4_72), inference(backward_demodulation, [], [f363, f912])).
fof(f6787, plain, (~ spl4_78 | ~ spl4_86 | ~ spl4_95 | spl4_177), inference(avatar_contradiction_clause, [], [f6786])).
fof(f6786, plain, ($false | (~ spl4_78 | ~ spl4_86 | ~ spl4_95 | spl4_177)), inference(subsumption_resolution, [], [f6780, f971])).
fof(f6780, plain, (~ (e0 = op(e1, e2)) | (~ spl4_78 | ~ spl4_95 | spl4_177)), inference(backward_demodulation, [], [f6665, f937])).
fof(f6665, plain, (~ (e0 = op(e1, op(e1, e4))) | (~ spl4_95 | spl4_177)), inference(forward_demodulation, [], [f1468, f1008])).
fof(f1468, plain, (~ (e0 = op(e1, op(e1, op(e1, e1)))) | spl4_177), inference(avatar_component_clause, [], [f1466])).
fof(f1466, plain, (spl4_177 <=> (e0 = op(e1, op(e1, op(e1, e1))))), introduced(avatar_definition, [new_symbols(naming, [spl4_177])])).
fof(f6769, plain, (~ spl4_57 | ~ spl4_82), inference(avatar_split_clause, [], [f6763, f952, f847])).
fof(f6763, plain, (~ (e1 = op(e2, e3)) | ~ spl4_82), inference(backward_demodulation, [], [f390, f954])).
fof(f6748, plain, (spl4_82 | ~ spl4_99 | ~ spl4_117), inference(avatar_split_clause, [], [f6744, f1099, f1023, f952])).
fof(f6744, plain, ((e1 = op(e1, e3)) | (~ spl4_99 | ~ spl4_117)), inference(backward_demodulation, [], [f5717, f1025])).
fof(f5717, plain, ((e1 = op(e1, op(e1, e0))) | ~ spl4_117), inference(backward_demodulation, [], [f336, f1101])).
fof(f6720, plain, (~ spl4_88 | ~ spl4_113), inference(avatar_split_clause, [], [f6719, f1082, f977])).
fof(f6719, plain, (~ (e2 = op(e1, e2)) | ~ spl4_113), inference(forward_demodulation, [], [f376, f1084])).
fof(f6716, plain, (~ spl4_85 | ~ spl4_95), inference(avatar_split_clause, [], [f6715, f1006, f964])).
fof(f6715, plain, (~ (e4 = op(e1, e3)) | ~ spl4_95), inference(forward_demodulation, [], [f421, f1008])).
fof(f6713, plain, (~ spl4_80 | ~ spl4_105), inference(avatar_split_clause, [], [f6712, f1048, f943])).
fof(f6712, plain, (~ (e4 = op(e1, e4)) | ~ spl4_105), inference(forward_demodulation, [], [f396, f1050])).
fof(f6704, plain, (~ spl4_55 | ~ spl4_105), inference(avatar_split_clause, [], [f6703, f1048, f838])).
fof(f6703, plain, (~ (e4 = op(e2, e4)) | ~ spl4_105), inference(forward_demodulation, [], [f397, f1050])).
fof(f6678, plain, (~ spl4_26 | ~ spl4_33 | ~ spl4_40 | spl4_159), inference(avatar_contradiction_clause, [], [f6677])).
fof(f6677, plain, ($false | (~ spl4_26 | ~ spl4_33 | ~ spl4_40 | spl4_159)), inference(subsumption_resolution, [], [f6676, f719])).
fof(f6676, plain, (~ (e0 = op(e3, e4)) | (~ spl4_33 | ~ spl4_40 | spl4_159)), inference(forward_demodulation, [], [f6675, f777])).
fof(f6675, plain, (~ (e0 = op(e3, op(e3, e2))) | (~ spl4_33 | spl4_159)), inference(forward_demodulation, [], [f1370, f748])).
fof(f6647, plain, (~ spl4_2 | ~ spl4_11 | ~ spl4_18 | spl4_161), inference(avatar_contradiction_clause, [], [f6646])).
fof(f6646, plain, ($false | (~ spl4_2 | ~ spl4_11 | ~ spl4_18 | spl4_161)), inference(subsumption_resolution, [], [f6645, f656])).
fof(f6645, plain, (~ (e0 = op(e4, e2)) | (~ spl4_2 | ~ spl4_18 | spl4_161)), inference(forward_demodulation, [], [f6644, f685])).
fof(f6644, plain, (~ (e0 = op(e4, op(e4, e1))) | (~ spl4_2 | spl4_161)), inference(forward_demodulation, [], [f1379, f618])).
fof(f1379, plain, (~ (e0 = op(e4, op(e4, op(e4, e4)))) | spl4_161), inference(avatar_component_clause, [], [f1377])).
fof(f1377, plain, (spl4_161 <=> (e0 = op(e4, op(e4, op(e4, e4))))), introduced(avatar_definition, [new_symbols(naming, [spl4_161])])).
fof(f6643, plain, (~ spl4_2 | ~ spl4_18 | spl4_139), inference(avatar_contradiction_clause, [], [f6642])).
fof(f6642, plain, ($false | (~ spl4_2 | ~ spl4_18 | spl4_139)), inference(subsumption_resolution, [], [f6638, f685])).
fof(f6638, plain, (~ (e2 = op(e4, e1)) | (~ spl4_2 | spl4_139)), inference(backward_demodulation, [], [f1269, f618])).
fof(f6619, plain, (~ spl4_1 | ~ spl4_26), inference(avatar_split_clause, [], [f6614, f717, f612])).
fof(f6614, plain, (~ (e0 = op(e4, e4)) | ~ spl4_26), inference(backward_demodulation, [], [f405, f719])).
fof(f6584, plain, (spl4_11 | ~ spl4_75 | ~ spl4_113), inference(avatar_split_clause, [], [f6581, f1082, f922, f654])).
fof(f6581, plain, ((e0 = op(e4, e2)) | (~ spl4_75 | ~ spl4_113)), inference(backward_demodulation, [], [f6564, f924])).
fof(f6564, plain, ((e0 = op(op(e2, e0), e2)) | ~ spl4_113), inference(backward_demodulation, [], [f333, f1084])).
fof(f6578, plain, (spl4_26 | ~ spl4_24 | ~ spl4_105), inference(avatar_split_clause, [], [f6575, f1048, f708, f717])).
fof(f6575, plain, ((e0 = op(e3, e4)) | (~ spl4_24 | ~ spl4_105)), inference(backward_demodulation, [], [f6495, f1050])).
fof(f6495, plain, ((e0 = op(e3, op(e0, e4))) | ~ spl4_24), inference(forward_demodulation, [], [f335, f710])).
fof(f6571, plain, (~ spl4_103 | ~ spl4_113), inference(avatar_split_clause, [], [f6567, f1082, f1040])).
fof(f6567, plain, (~ (e2 = op(e0, e4)) | ~ spl4_113), inference(backward_demodulation, [], [f414, f1084])).
fof(f6554, plain, (~ spl4_46 | ~ spl4_121), inference(avatar_split_clause, [], [f6543, f1116, f801])).
fof(f6543, plain, (~ (e0 = op(e3, e0)) | ~ spl4_121), inference(backward_demodulation, [], [f358, f1118])).
fof(f6540, plain, (~ spl4_33 | ~ spl4_40 | spl4_132), inference(avatar_contradiction_clause, [], [f6539])).
fof(f6539, plain, ($false | (~ spl4_33 | ~ spl4_40 | spl4_132)), inference(subsumption_resolution, [], [f6538, f777])).
fof(f6538, plain, (~ (e4 = op(e3, e2)) | (~ spl4_33 | spl4_132)), inference(forward_demodulation, [], [f1238, f748])).
fof(f6526, plain, (~ spl4_33 | ~ spl4_135 | ~ spl4_183), inference(avatar_contradiction_clause, [], [f6525])).
fof(f6525, plain, ($false | (~ spl4_33 | ~ spl4_135 | ~ spl4_183)), inference(subsumption_resolution, [], [f6524, f463])).
fof(f6524, plain, ((e2 = e3) | (~ spl4_33 | ~ spl4_135 | ~ spl4_183)), inference(forward_demodulation, [], [f6523, f748])).
fof(f6523, plain, ((e3 = op(e3, e3)) | (~ spl4_135 | ~ spl4_183)), inference(forward_demodulation, [], [f1514, f1250])).
fof(f1514, plain, ((e3 = op(op(e4, op(e4, e4)), op(e4, op(e4, e4)))) | ~ spl4_183), inference(avatar_component_clause, [], [f1513])).
fof(f6482, plain, (spl4_57 | ~ spl4_18 | ~ spl4_79), inference(avatar_split_clause, [], [f6445, f939, f683, f847])).
fof(f6445, plain, ((e1 = op(e2, e3)) | (~ spl4_18 | ~ spl4_79)), inference(backward_demodulation, [], [f6317, f685])).
fof(f6317, plain, ((e1 = op(op(e4, e1), e3)) | ~ spl4_79), inference(backward_demodulation, [], [f340, f941])).
fof(f6472, plain, (~ spl4_8 | ~ spl4_135 | spl4_143), inference(avatar_split_clause, [], [f6232, f1287, f1249, f641])).
fof(f6232, plain, (~ (e2 = op(e4, e3)) | (~ spl4_135 | spl4_143)), inference(backward_demodulation, [], [f1289, f1250])).
fof(f6469, plain, (~ spl4_1 | ~ spl4_5), inference(avatar_contradiction_clause, [], [f6468])).
fof(f6468, plain, ($false | (~ spl4_1 | ~ spl4_5)), inference(subsumption_resolution, [], [f6467, f459])).
fof(f459, plain, ~ (e0 = e4), inference(cnf_transformation, [], [f5])).
fof(f6467, plain, ((e0 = e4) | (~ spl4_1 | ~ spl4_5)), inference(backward_demodulation, [], [f630, f614])).
fof(f6463, plain, (spl4_4 | ~ spl4_5 | ~ spl4_135), inference(avatar_split_clause, [], [f6461, f1249, f628, f624])).
fof(f6461, plain, ((e3 = op(e4, e4)) | (~ spl4_5 | ~ spl4_135)), inference(backward_demodulation, [], [f1250, f630])).
fof(f6456, plain, (~ spl4_13 | ~ spl4_53 | spl4_63), inference(avatar_contradiction_clause, [], [f6455])).
fof(f6455, plain, ($false | (~ spl4_13 | ~ spl4_53 | spl4_63)), inference(subsumption_resolution, [], [f6453, f873])).
fof(f873, plain, (~ (e2 = op(e2, e2)) | spl4_63), inference(avatar_component_clause, [], [f872])).
fof(f6453, plain, ((e2 = op(e2, e2)) | (~ spl4_13 | ~ spl4_53)), inference(backward_demodulation, [], [f6378, f664])).
fof(f6378, plain, ((e2 = op(op(e4, e2), e2)) | ~ spl4_53), inference(backward_demodulation, [], [f345, f832])).
fof(f6438, plain, (~ spl4_4 | ~ spl4_24), inference(avatar_split_clause, [], [f6434, f708, f624])).
fof(f6434, plain, (~ (e3 = op(e4, e4)) | ~ spl4_24), inference(backward_demodulation, [], [f449, f710])).
fof(f6431, plain, (~ spl4_7 | ~ spl4_27 | spl4_94), inference(avatar_contradiction_clause, [], [f6430])).
fof(f6430, plain, ($false | (~ spl4_7 | ~ spl4_27 | spl4_94)), inference(subsumption_resolution, [], [f6429, f1003])).
fof(f1003, plain, (~ (e3 = op(e1, e1)) | spl4_94), inference(avatar_component_clause, [], [f1002])).
fof(f6429, plain, ((e3 = op(e1, e1)) | (~ spl4_7 | ~ spl4_27)), inference(backward_demodulation, [], [f6202, f723])).
fof(f6202, plain, ((e3 = op(e1, op(e3, e4))) | ~ spl4_7), inference(forward_demodulation, [], [f350, f639])).
fof(f350, plain, (e3 = op(op(e4, e3), op(e3, e4))), inference(cnf_transformation, [], [f3])).
fof(f6410, plain, (spl4_4 | ~ spl4_40 | ~ spl4_60), inference(avatar_split_clause, [], [f6405, f859, f775, f624])).
fof(f6405, plain, ((e3 = op(e4, e4)) | (~ spl4_40 | ~ spl4_60)), inference(backward_demodulation, [], [f6367, f777])).
fof(f6367, plain, ((e3 = op(e4, op(e3, e2))) | ~ spl4_60), inference(backward_demodulation, [], [f348, f861])).
fof(f6409, plain, (spl4_3 | ~ spl4_40 | ~ spl4_60), inference(avatar_contradiction_clause, [], [f6408])).
fof(f6408, plain, ($false | (spl4_3 | ~ spl4_40 | ~ spl4_60)), inference(subsumption_resolution, [], [f6404, f621])).
fof(f621, plain, (~ (e2 = op(e4, e4)) | spl4_3), inference(avatar_component_clause, [], [f620])).
fof(f6404, plain, ((e2 = op(e4, e4)) | (~ spl4_40 | ~ spl4_60)), inference(backward_demodulation, [], [f6366, f777])).
fof(f6366, plain, ((e2 = op(op(e3, e2), e4)) | ~ spl4_60), inference(backward_demodulation, [], [f344, f861])).
fof(f6407, plain, (~ spl4_35 | ~ spl4_40), inference(avatar_split_clause, [], [f6402, f775, f754])).
fof(f6402, plain, (~ (e4 = op(e3, e3)) | ~ spl4_40), inference(backward_demodulation, [], [f443, f777])).
fof(f6395, plain, (~ spl4_36 | ~ spl4_46), inference(avatar_split_clause, [], [f6386, f801, f759])).
fof(f6386, plain, (~ (e0 = op(e3, e2)) | ~ spl4_46), inference(backward_demodulation, [], [f437, f803])).
fof(f6352, plain, (~ spl4_16 | ~ spl4_66), inference(avatar_split_clause, [], [f6342, f885, f675])).
fof(f6342, plain, (~ (e0 = op(e4, e1)) | ~ spl4_66), inference(backward_demodulation, [], [f374, f887])).
fof(f6338, plain, (~ spl4_72 | spl4_96 | ~ spl4_111), inference(avatar_contradiction_clause, [], [f6337])).
fof(f6337, plain, ($false | (~ spl4_72 | spl4_96 | ~ spl4_111)), inference(subsumption_resolution, [], [f6333, f1012])).
fof(f6333, plain, ((e0 = op(e1, e0)) | (~ spl4_72 | ~ spl4_111)), inference(backward_demodulation, [], [f6253, f912])).
fof(f6253, plain, ((e0 = op(op(e2, e0), e0)) | ~ spl4_111), inference(backward_demodulation, [], [f333, f1076])).
fof(f6316, plain, (~ spl4_77 | ~ spl4_87), inference(avatar_split_clause, [], [f6311, f973, f931])).
fof(f6311, plain, (~ (e1 = op(e1, e4)) | ~ spl4_87), inference(backward_demodulation, [], [f424, f975])).
fof(f6315, plain, (~ spl4_62 | ~ spl4_87), inference(avatar_split_clause, [], [f6308, f973, f868])).
fof(f6308, plain, (~ (e1 = op(e2, e2)) | ~ spl4_87), inference(backward_demodulation, [], [f380, f975])).
fof(f6304, plain, (~ spl4_70 | ~ spl4_95), inference(avatar_split_clause, [], [f6296, f1006, f901])).
fof(f6296, plain, (~ (e4 = op(e2, e1)) | ~ spl4_95), inference(backward_demodulation, [], [f370, f1008])).
fof(f6294, plain, (spl4_87 | ~ spl4_98 | ~ spl4_117), inference(avatar_split_clause, [], [f6290, f1099, f1019, f973])).
fof(f6290, plain, ((e1 = op(e1, e2)) | (~ spl4_98 | ~ spl4_117)), inference(backward_demodulation, [], [f5717, f1021])).
fof(f6292, plain, (~ spl4_93 | ~ spl4_98), inference(avatar_split_clause, [], [f6284, f1019, f998])).
fof(f6284, plain, (~ (e2 = op(e1, e1)) | ~ spl4_98), inference(backward_demodulation, [], [f416, f1021])).
fof(f6279, plain, (~ spl4_62 | ~ spl4_103 | ~ spl4_125 | spl4_164), inference(avatar_split_clause, [], [f6273, f1396, f1132, f1040, f868])).
fof(f6273, plain, (~ (e1 = op(e2, e2)) | (~ spl4_103 | ~ spl4_125 | spl4_164)), inference(backward_demodulation, [], [f6246, f1042])).
fof(f6246, plain, (~ (e1 = op(op(e0, e4), op(e0, e4))) | (~ spl4_125 | spl4_164)), inference(backward_demodulation, [], [f1398, f1134])).
fof(f6214, plain, (~ spl4_33 | ~ spl4_148 | spl4_176), inference(avatar_split_clause, [], [f6213, f1462, f1310, f746])).
fof(f1462, plain, (spl4_176 <=> (e2 = op(op(e1, op(e1, e1)), op(e1, op(e1, e1))))), introduced(avatar_definition, [new_symbols(naming, [spl4_176])])).
fof(f6213, plain, (~ (e2 = op(e3, e3)) | (~ spl4_148 | spl4_176)), inference(forward_demodulation, [], [f1464, f1311])).
fof(f1464, plain, (~ (e2 = op(op(e1, op(e1, e1)), op(e1, op(e1, e1)))) | spl4_176), inference(avatar_component_clause, [], [f1462])).
fof(f6207, plain, (spl4_47 | ~ spl4_44 | ~ spl4_81), inference(avatar_split_clause, [], [f6118, f948, f792, f805])).
fof(f6118, plain, ((e1 = op(e3, e0)) | (~ spl4_44 | ~ spl4_81)), inference(backward_demodulation, [], [f6037, f794])).
fof(f6037, plain, ((e1 = op(op(e3, e1), e0)) | ~ spl4_81), inference(backward_demodulation, [], [f339, f950])).
fof(f6197, plain, (~ spl4_19 | ~ spl4_44), inference(avatar_split_clause, [], [f6196, f792, f687])).
fof(f6196, plain, (~ (e3 = op(e4, e1)) | ~ spl4_44), inference(forward_demodulation, [], [f375, f794])).
fof(f6193, plain, (~ spl4_62 | ~ spl4_70 | spl4_138), inference(avatar_contradiction_clause, [], [f6192])).
fof(f6192, plain, ($false | (~ spl4_62 | ~ spl4_70 | spl4_138)), inference(subsumption_resolution, [], [f6191, f903])).
fof(f6191, plain, (~ (e4 = op(e2, e1)) | (~ spl4_62 | spl4_138)), inference(forward_demodulation, [], [f1264, f870])).
fof(f1264, plain, (~ (e4 = op(e2, op(e2, e2))) | spl4_138), inference(avatar_component_clause, [], [f1262])).
fof(f1262, plain, (spl4_138 <=> (e4 = op(e2, op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl4_138])])).
fof(f6190, plain, (~ spl4_28 | ~ spl4_35 | spl4_141), inference(avatar_contradiction_clause, [], [f6189])).
fof(f6189, plain, ($false | (~ spl4_28 | ~ spl4_35 | spl4_141)), inference(subsumption_resolution, [], [f6188, f727])).
fof(f6188, plain, (~ (e2 = op(e3, e4)) | (~ spl4_35 | spl4_141)), inference(forward_demodulation, [], [f1279, f756])).
fof(f6187, plain, (~ spl4_4 | ~ spl4_7 | spl4_147), inference(avatar_contradiction_clause, [], [f6186])).
fof(f6186, plain, ($false | (~ spl4_4 | ~ spl4_7 | spl4_147)), inference(subsumption_resolution, [], [f6185, f639])).
fof(f6185, plain, (~ (e1 = op(e4, e3)) | (~ spl4_4 | spl4_147)), inference(forward_demodulation, [], [f1307, f626])).
fof(f6139, plain, (~ spl4_28 | ~ spl4_35 | ~ spl4_62 | spl4_158), inference(avatar_contradiction_clause, [], [f6138])).
fof(f6138, plain, ($false | (~ spl4_28 | ~ spl4_35 | ~ spl4_62 | spl4_158)), inference(subsumption_resolution, [], [f6134, f870])).
fof(f6134, plain, (~ (e1 = op(e2, e2)) | (~ spl4_28 | ~ spl4_35 | spl4_158)), inference(backward_demodulation, [], [f5600, f727])).
fof(f5600, plain, (~ (e1 = op(op(e3, e4), op(e3, e4))) | (~ spl4_35 | spl4_158)), inference(backward_demodulation, [], [f1366, f756])).
fof(f6120, plain, (~ spl4_39 | ~ spl4_44), inference(avatar_split_clause, [], [f6115, f792, f771])).
fof(f6115, plain, (~ (e3 = op(e3, e2)) | ~ spl4_44), inference(backward_demodulation, [], [f440, f794])).
fof(f6114, plain, (spl4_44 | ~ spl4_47 | ~ spl4_109), inference(avatar_split_clause, [], [f6112, f1065, f805, f792])).
fof(f6112, plain, ((e3 = op(e3, e1)) | (~ spl4_47 | ~ spl4_109)), inference(backward_demodulation, [], [f5995, f807])).
fof(f6094, plain, (~ spl4_4 | ~ spl4_62 | ~ spl4_70 | spl4_182), inference(avatar_contradiction_clause, [], [f6093])).
fof(f6093, plain, ($false | (~ spl4_4 | ~ spl4_62 | ~ spl4_70 | spl4_182)), inference(subsumption_resolution, [], [f6092, f626])).
fof(f6092, plain, (~ (e3 = op(e4, e4)) | (~ spl4_62 | ~ spl4_70 | spl4_182)), inference(forward_demodulation, [], [f6082, f903])).
fof(f6082, plain, (~ (e3 = op(op(e2, e1), op(e2, e1))) | (~ spl4_62 | spl4_182)), inference(backward_demodulation, [], [f1510, f870])).
fof(f1510, plain, (~ (e3 = op(op(e2, op(e2, e2)), op(e2, op(e2, e2)))) | spl4_182), inference(avatar_component_clause, [], [f1508])).
fof(f1508, plain, (spl4_182 <=> (e3 = op(op(e2, op(e2, e2)), op(e2, op(e2, e2))))), introduced(avatar_definition, [new_symbols(naming, [spl4_182])])).
fof(f6091, plain, (~ spl4_51 | ~ spl4_62 | ~ spl4_70 | spl4_163), inference(avatar_contradiction_clause, [], [f6090])).
fof(f6090, plain, ($false | (~ spl4_51 | ~ spl4_62 | ~ spl4_70 | spl4_163)), inference(subsumption_resolution, [], [f6089, f824])).
fof(f6089, plain, (~ (e0 = op(e2, e4)) | (~ spl4_62 | ~ spl4_70 | spl4_163)), inference(forward_demodulation, [], [f6081, f903])).
fof(f6081, plain, (~ (e0 = op(e2, op(e2, e1))) | (~ spl4_62 | spl4_163)), inference(backward_demodulation, [], [f1388, f870])).
fof(f1388, plain, (~ (e0 = op(e2, op(e2, op(e2, e2)))) | spl4_163), inference(avatar_component_clause, [], [f1386])).
fof(f1386, plain, (spl4_163 <=> (e0 = op(e2, op(e2, op(e2, e2))))), introduced(avatar_definition, [new_symbols(naming, [spl4_163])])).
fof(f6085, plain, (~ spl4_57 | ~ spl4_62), inference(avatar_split_clause, [], [f6075, f868, f847])).
fof(f6075, plain, (~ (e1 = op(e2, e3)) | ~ spl4_62), inference(backward_demodulation, [], [f433, f870])).
fof(f6071, plain, (spl4_28 | ~ spl4_70 | ~ spl4_89), inference(avatar_split_clause, [], [f6068, f981, f901, f725])).
fof(f6068, plain, ((e2 = op(e3, e4)) | (~ spl4_70 | ~ spl4_89)), inference(backward_demodulation, [], [f6020, f903])).
fof(f6020, plain, ((e2 = op(e3, op(e2, e1))) | ~ spl4_89), inference(backward_demodulation, [], [f342, f983])).
fof(f6070, plain, (spl4_7 | ~ spl4_70 | ~ spl4_89), inference(avatar_split_clause, [], [f6067, f981, f901, f637])).
fof(f6067, plain, ((e1 = op(e4, e3)) | (~ spl4_70 | ~ spl4_89)), inference(backward_demodulation, [], [f6019, f903])).
fof(f6019, plain, ((e1 = op(op(e2, e1), e3)) | ~ spl4_89), inference(backward_demodulation, [], [f338, f983])).
fof(f6034, plain, (~ spl4_84 | ~ spl4_89), inference(avatar_split_clause, [], [f6024, f981, f960])).
fof(f6024, plain, (~ (e3 = op(e1, e3)) | ~ spl4_89), inference(backward_demodulation, [], [f423, f983])).
fof(f6033, plain, (~ spl4_39 | ~ spl4_89), inference(avatar_split_clause, [], [f6022, f981, f771])).
fof(f6022, plain, (~ (e3 = op(e3, e2)) | ~ spl4_89), inference(backward_demodulation, [], [f381, f983])).
fof(f6001, plain, (~ spl4_84 | ~ spl4_109), inference(avatar_split_clause, [], [f5996, f1065, f960])).
fof(f5996, plain, (~ (e3 = op(e1, e3)) | ~ spl4_109), inference(backward_demodulation, [], [f386, f1067])).
fof(f5992, plain, (~ spl4_63 | ~ spl4_113), inference(avatar_split_clause, [], [f5985, f1082, f872])).
fof(f5985, plain, (~ (e2 = op(e2, e2)) | ~ spl4_113), inference(backward_demodulation, [], [f377, f1084])).
fof(f5973, plain, (~ spl4_96 | ~ spl4_121), inference(avatar_split_clause, [], [f5958, f1116, f1011])).
fof(f5958, plain, (~ (e0 = op(e1, e0)) | ~ spl4_121), inference(backward_demodulation, [], [f356, f1118])).
fof(f5950, plain, (~ spl4_104 | ~ spl4_4), inference(avatar_split_clause, [], [f5949, f624, f1044])).
fof(f5949, plain, (~ (e3 = op(e0, e4)) | ~ spl4_4), inference(forward_demodulation, [], [f399, f626])).
fof(f5928, plain, (spl4_35 | ~ spl4_4), inference(avatar_split_clause, [], [f5896, f624, f754])).
fof(f5896, plain, ((e4 = op(e3, e3)) | ~ spl4_4), inference(backward_demodulation, [], [f355, f626])).
fof(f355, plain, (e4 = op(op(e4, e4), op(e4, e4))), inference(cnf_transformation, [], [f3])).
fof(f5912, plain, (~ spl4_4 | ~ spl4_8 | ~ spl4_12 | spl4_134), inference(avatar_contradiction_clause, [], [f5911])).
fof(f5911, plain, ($false | (~ spl4_4 | ~ spl4_8 | ~ spl4_12 | spl4_134)), inference(subsumption_resolution, [], [f5910, f660])).
fof(f5910, plain, (~ (e1 = op(e4, e2)) | (~ spl4_4 | ~ spl4_8 | spl4_134)), inference(forward_demodulation, [], [f5899, f643])).
fof(f5899, plain, (~ (e1 = op(e4, op(e4, e3))) | (~ spl4_4 | spl4_134)), inference(backward_demodulation, [], [f1247, f626])).
fof(f5895, plain, (~ spl4_12 | ~ spl4_51 | ~ spl4_117), inference(avatar_contradiction_clause, [], [f5894])).
fof(f5894, plain, ($false | (~ spl4_12 | ~ spl4_51 | ~ spl4_117)), inference(subsumption_resolution, [], [f5893, f462])).
fof(f5893, plain, ((e1 = e4) | (~ spl4_12 | ~ spl4_51 | ~ spl4_117)), inference(forward_demodulation, [], [f5889, f1101])).
fof(f5889, plain, ((e4 = op(e0, e1)) | (~ spl4_12 | ~ spl4_51)), inference(backward_demodulation, [], [f5842, f660])).
fof(f5842, plain, ((e4 = op(e0, op(e4, e2))) | ~ spl4_51), inference(backward_demodulation, [], [f353, f824])).
fof(f5892, plain, (~ spl4_12 | ~ spl4_51 | ~ spl4_96), inference(avatar_contradiction_clause, [], [f5891])).
fof(f5891, plain, ($false | (~ spl4_12 | ~ spl4_51 | ~ spl4_96)), inference(subsumption_resolution, [], [f5890, f457])).
fof(f5890, plain, ((e0 = e2) | (~ spl4_12 | ~ spl4_51 | ~ spl4_96)), inference(forward_demodulation, [], [f5888, f1013])).
fof(f5888, plain, ((e2 = op(e1, e0)) | (~ spl4_12 | ~ spl4_51)), inference(backward_demodulation, [], [f5841, f660])).
fof(f5841, plain, ((e2 = op(op(e4, e2), e0)) | ~ spl4_51), inference(backward_demodulation, [], [f345, f824])).
fof(f5883, plain, (~ spl4_16 | ~ spl4_80 | spl4_102), inference(avatar_contradiction_clause, [], [f5882])).
fof(f5882, plain, ($false | (~ spl4_16 | ~ spl4_80 | spl4_102)), inference(subsumption_resolution, [], [f5879, f1037])).
fof(f5879, plain, ((e1 = op(e0, e4)) | (~ spl4_16 | ~ spl4_80)), inference(backward_demodulation, [], [f5029, f677])).
fof(f5029, plain, ((e1 = op(op(e4, e1), e4)) | ~ spl4_80), inference(backward_demodulation, [], [f340, f945])).
fof(f5881, plain, (~ spl4_1 | ~ spl4_16), inference(avatar_split_clause, [], [f5874, f675, f612])).
fof(f5874, plain, (~ (e0 = op(e4, e4)) | ~ spl4_16), inference(backward_demodulation, [], [f452, f677])).
fof(f5867, plain, (~ spl4_15 | ~ spl4_25), inference(avatar_split_clause, [], [f5863, f712, f670])).
fof(f5863, plain, (~ (e4 = op(e4, e2)) | ~ spl4_25), inference(backward_demodulation, [], [f447, f714])).
fof(f5861, plain, (~ spl4_31 | ~ spl4_35), inference(avatar_contradiction_clause, [], [f5860])).
fof(f5860, plain, ($false | (~ spl4_31 | ~ spl4_35)), inference(subsumption_resolution, [], [f5859, f459])).
fof(f5859, plain, ((e0 = e4) | (~ spl4_31 | ~ spl4_35)), inference(backward_demodulation, [], [f756, f740])).
fof(f5847, plain, (~ spl4_1 | ~ spl4_51), inference(avatar_split_clause, [], [f5843, f822, f612])).
fof(f5843, plain, (~ (e0 = op(e4, e4)) | ~ spl4_51), inference(backward_demodulation, [], [f404, f824])).
fof(f5840, plain, (~ spl4_69 | spl4_83 | ~ spl4_87), inference(avatar_contradiction_clause, [], [f5839])).
fof(f5839, plain, ($false | (~ spl4_69 | spl4_83 | ~ spl4_87)), inference(subsumption_resolution, [], [f5834, f957])).
fof(f957, plain, (~ (e2 = op(e1, e3)) | spl4_83), inference(avatar_component_clause, [], [f956])).
fof(f5834, plain, ((e2 = op(e1, e3)) | (~ spl4_69 | ~ spl4_87)), inference(backward_demodulation, [], [f5791, f899])).
fof(f5791, plain, ((e2 = op(e1, op(e2, e1))) | ~ spl4_87), inference(backward_demodulation, [], [f342, f975])).
fof(f5836, plain, (~ spl4_54 | ~ spl4_69), inference(avatar_split_clause, [], [f5831, f897, f834])).
fof(f5831, plain, (~ (e3 = op(e2, e4)) | ~ spl4_69), inference(backward_demodulation, [], [f432, f899])).
fof(f5835, plain, (~ spl4_19 | ~ spl4_69), inference(avatar_split_clause, [], [f5830, f897, f687])).
fof(f5830, plain, (~ (e3 = op(e4, e1)) | ~ spl4_69), inference(backward_demodulation, [], [f374, f899])).
fof(f5828, plain, (~ spl4_72 | ~ spl4_75), inference(avatar_contradiction_clause, [], [f5827])).
fof(f5827, plain, ($false | (~ spl4_72 | ~ spl4_75)), inference(subsumption_resolution, [], [f5826, f462])).
fof(f5826, plain, ((e1 = e4) | (~ spl4_72 | ~ spl4_75)), inference(backward_demodulation, [], [f924, f912])).
fof(f5803, plain, (~ spl4_12 | ~ spl4_87), inference(avatar_split_clause, [], [f5792, f973, f658])).
fof(f5792, plain, (~ (e1 = op(e4, e2)) | ~ spl4_87), inference(backward_demodulation, [], [f382, f975])).
fof(f5789, plain, (~ spl4_88 | ~ spl4_93), inference(avatar_split_clause, [], [f5774, f998, f977])).
fof(f5774, plain, (~ (e2 = op(e1, e2)) | ~ spl4_93), inference(backward_demodulation, [], [f420, f1000])).
fof(f5769, plain, (~ spl4_96 | ~ spl4_117), inference(avatar_contradiction_clause, [], [f5768])).
fof(f5768, plain, ($false | (~ spl4_96 | ~ spl4_117)), inference(subsumption_resolution, [], [f5767, f456])).
fof(f456, plain, ~ (e0 = e1), inference(cnf_transformation, [], [f5])).
fof(f5767, plain, ((e0 = e1) | (~ spl4_96 | ~ spl4_117)), inference(forward_demodulation, [], [f5762, f1013])).
fof(f5762, plain, ((e1 = op(e1, e0)) | (~ spl4_96 | ~ spl4_117)), inference(backward_demodulation, [], [f5717, f1013])).
fof(f5764, plain, (~ spl4_81 | ~ spl4_96), inference(avatar_split_clause, [], [f5756, f1011, f948])).
fof(f5756, plain, (~ (e0 = op(e1, e3)) | ~ spl4_96), inference(backward_demodulation, [], [f418, f1013])).
fof(f5750, plain, (spl4_51 | ~ spl4_48 | ~ spl4_110), inference(avatar_split_clause, [], [f5744, f1069, f809, f822])).
fof(f5744, plain, ((e0 = op(e2, e4)) | (~ spl4_48 | ~ spl4_110)), inference(backward_demodulation, [], [f5688, f1071])).
fof(f5688, plain, ((e0 = op(e2, op(e0, e3))) | ~ spl4_48), inference(forward_demodulation, [], [f334, f811])).
fof(f5749, plain, (spl4_14 | ~ spl4_48 | ~ spl4_110), inference(avatar_contradiction_clause, [], [f5748])).
fof(f5748, plain, ($false | (spl4_14 | ~ spl4_48 | ~ spl4_110)), inference(subsumption_resolution, [], [f5743, f667])).
fof(f667, plain, (~ (e3 = op(e4, e2)) | spl4_14), inference(avatar_component_clause, [], [f666])).
fof(f5743, plain, ((e3 = op(e4, e2)) | (~ spl4_48 | ~ spl4_110)), inference(backward_demodulation, [], [f5687, f1071])).
fof(f5687, plain, ((e3 = op(op(e0, e3), e2)) | ~ spl4_48), inference(forward_demodulation, [], [f346, f811])).
fof(f5741, plain, (~ spl4_111 | ~ spl4_115), inference(avatar_contradiction_clause, [], [f5740])).
fof(f5740, plain, ($false | (~ spl4_111 | ~ spl4_115)), inference(subsumption_resolution, [], [f5739, f459])).
fof(f5739, plain, ((e0 = e4) | (~ spl4_111 | ~ spl4_115)), inference(backward_demodulation, [], [f1092, f1076])).
fof(f5727, plain, (~ spl4_92 | ~ spl4_117), inference(avatar_split_clause, [], [f5718, f1099, f994])).
fof(f5718, plain, (~ (e1 = op(e1, e1)) | ~ spl4_117), inference(backward_demodulation, [], [f366, f1101])).
fof(f5714, plain, (~ spl4_109 | ~ spl4_124), inference(avatar_split_clause, [], [f5706, f1128, f1065])).
fof(f5706, plain, (~ (e3 = op(e0, e3)) | ~ spl4_124), inference(backward_demodulation, [], [f408, f1130])).
fof(f5686, plain, (~ spl4_110 | ~ spl4_35), inference(avatar_split_clause, [], [f5685, f754, f1069])).
fof(f5685, plain, (~ (e4 = op(e0, e3)) | ~ spl4_35), inference(forward_demodulation, [], [f388, f756])).
fof(f5680, plain, (~ spl4_93 | ~ spl4_27 | ~ spl4_35 | spl4_174), inference(avatar_split_clause, [], [f5614, f1452, f754, f721, f998])).
fof(f1452, plain, (spl4_174 <=> (e2 = op(op(e3, op(e3, e3)), op(e3, op(e3, e3))))), introduced(avatar_definition, [new_symbols(naming, [spl4_174])])).
fof(f5614, plain, (~ (e2 = op(e1, e1)) | (~ spl4_27 | ~ spl4_35 | spl4_174)), inference(forward_demodulation, [], [f5602, f723])).
fof(f5602, plain, (~ (e2 = op(op(e3, e4), op(e3, e4))) | (~ spl4_35 | spl4_174)), inference(backward_demodulation, [], [f1454, f756])).
fof(f1454, plain, (~ (e2 = op(op(e3, op(e3, e3)), op(e3, op(e3, e3)))) | spl4_174), inference(avatar_component_clause, [], [f1452])).
fof(f5678, plain, (spl4_90 | ~ spl4_8 | ~ spl4_27), inference(avatar_split_clause, [], [f5634, f721, f641, f985])).
fof(f5634, plain, ((e4 = op(e1, e2)) | (~ spl4_8 | ~ spl4_27)), inference(backward_demodulation, [], [f5393, f643])).
fof(f5393, plain, ((e4 = op(e1, op(e4, e3))) | ~ spl4_27), inference(backward_demodulation, [], [f354, f723])).
fof(f5675, plain, (spl4_84 | ~ spl4_39 | ~ spl4_57), inference(avatar_split_clause, [], [f5585, f847, f771, f960])).
fof(f5585, plain, ((e3 = op(e1, e3)) | (~ spl4_39 | ~ spl4_57)), inference(backward_demodulation, [], [f5322, f773])).
fof(f5322, plain, ((e3 = op(e1, op(e3, e2))) | ~ spl4_57), inference(backward_demodulation, [], [f348, f849])).
fof(f5668, plain, (spl4_69 | ~ spl4_8 | ~ spl4_27), inference(avatar_split_clause, [], [f5633, f721, f641, f897])).
fof(f5633, plain, ((e3 = op(e2, e1)) | (~ spl4_8 | ~ spl4_27)), inference(backward_demodulation, [], [f5392, f643])).
fof(f5392, plain, ((e3 = op(op(e4, e3), e1)) | ~ spl4_27), inference(backward_demodulation, [], [f350, f723])).
fof(f5659, plain, (spl4_43 | ~ spl4_39 | ~ spl4_57), inference(avatar_split_clause, [], [f5584, f847, f771, f788])).
fof(f5584, plain, ((e2 = op(e3, e1)) | (~ spl4_39 | ~ spl4_57)), inference(backward_demodulation, [], [f5321, f773])).
fof(f5321, plain, ((e2 = op(op(e3, e2), e1)) | ~ spl4_57), inference(backward_demodulation, [], [f344, f849])).
fof(f5651, plain, (~ spl4_13 | ~ spl4_8), inference(avatar_split_clause, [], [f5650, f641, f662])).
fof(f5650, plain, (~ (e2 = op(e4, e2)) | ~ spl4_8), inference(forward_demodulation, [], [f453, f643])).
fof(f5621, plain, (~ spl4_14 | ~ spl4_19), inference(avatar_split_clause, [], [f5616, f687, f666])).
fof(f5616, plain, (~ (e3 = op(e4, e2)) | ~ spl4_19), inference(backward_demodulation, [], [f450, f689])).
fof(f5606, plain, (~ spl4_10 | ~ spl4_35), inference(avatar_split_clause, [], [f5596, f754, f649])).
fof(f5596, plain, (~ (e4 = op(e4, e3)) | ~ spl4_35), inference(backward_demodulation, [], [f395, f756])).
fof(f5605, plain, (~ spl4_1 | ~ spl4_35), inference(avatar_contradiction_clause, [], [f5604])).
fof(f5604, plain, ($false | (~ spl4_1 | ~ spl4_35)), inference(subsumption_resolution, [], [f5603, f458])).
fof(f458, plain, ~ (e0 = e3), inference(cnf_transformation, [], [f5])).
fof(f5603, plain, ((e0 = e3) | (~ spl4_1 | ~ spl4_35)), inference(forward_demodulation, [], [f5595, f614])).
fof(f5595, plain, ((e3 = op(e4, e4)) | ~ spl4_35), inference(backward_demodulation, [], [f349, f756])).
fof(f5540, plain, (~ spl4_84 | ~ spl4_99), inference(avatar_split_clause, [], [f5535, f1023, f960])).
fof(f5535, plain, (~ (e3 = op(e1, e3)) | ~ spl4_99), inference(backward_demodulation, [], [f418, f1025])).
fof(f5539, plain, (~ spl4_49 | ~ spl4_99), inference(avatar_split_clause, [], [f5533, f1023, f813])).
fof(f5533, plain, (~ (e3 = op(e3, e0)) | ~ spl4_99), inference(backward_demodulation, [], [f361, f1025])).
fof(f5528, plain, (spl4_86 | ~ spl4_22 | ~ spl4_103), inference(avatar_split_clause, [], [f5522, f1040, f700, f969])).
fof(f5522, plain, ((e0 = op(e1, e2)) | (~ spl4_22 | ~ spl4_103)), inference(backward_demodulation, [], [f4903, f1042])).
fof(f4903, plain, ((e0 = op(e1, op(e0, e4))) | ~ spl4_22), inference(forward_demodulation, [], [f335, f702])).
fof(f5488, plain, (spl4_98 | ~ spl4_71 | ~ spl4_112), inference(avatar_split_clause, [], [f5280, f1078, f906, f1019])).
fof(f5280, plain, ((e2 = op(e1, e0)) | (~ spl4_71 | ~ spl4_112)), inference(backward_demodulation, [], [f5202, f908])).
fof(f5202, plain, ((e2 = op(e1, op(e2, e0))) | ~ spl4_112), inference(backward_demodulation, [], [f341, f1080])).
fof(f5466, plain, (~ spl4_19 | ~ spl4_1 | ~ spl4_22 | spl4_151), inference(avatar_split_clause, [], [f5465, f1325, f700, f612, f687])).
fof(f5465, plain, (~ (e3 = op(e4, e1)) | (~ spl4_1 | ~ spl4_22 | spl4_151)), inference(forward_demodulation, [], [f5464, f702])).
fof(f5464, plain, (~ (e3 = op(e4, op(e4, e0))) | (~ spl4_1 | spl4_151)), inference(forward_demodulation, [], [f1327, f614])).
fof(f5432, plain, (~ spl4_1 | ~ spl4_3), inference(avatar_contradiction_clause, [], [f5431])).
fof(f5431, plain, ($false | (~ spl4_1 | ~ spl4_3)), inference(subsumption_resolution, [], [f5430, f457])).
fof(f5430, plain, ((e0 = e2) | (~ spl4_1 | ~ spl4_3)), inference(backward_demodulation, [], [f622, f614])).
fof(f5428, plain, (~ spl4_10 | ~ spl4_18 | ~ spl4_27), inference(avatar_contradiction_clause, [], [f5427])).
fof(f5427, plain, ($false | (~ spl4_10 | ~ spl4_18 | ~ spl4_27)), inference(subsumption_resolution, [], [f5426, f463])).
fof(f5426, plain, ((e2 = e3) | (~ spl4_10 | ~ spl4_18 | ~ spl4_27)), inference(forward_demodulation, [], [f5424, f685])).
fof(f5424, plain, ((e3 = op(e4, e1)) | (~ spl4_10 | ~ spl4_27)), inference(backward_demodulation, [], [f5392, f651])).
fof(f5416, plain, (~ spl4_9 | ~ spl4_14), inference(avatar_split_clause, [], [f5408, f666, f645])).
fof(f5408, plain, (~ (e3 = op(e4, e3)) | ~ spl4_14), inference(backward_demodulation, [], [f453, f668])).
fof(f5405, plain, (~ spl4_18 | spl4_52 | ~ spl4_80), inference(avatar_contradiction_clause, [], [f5404])).
fof(f5404, plain, ($false | (~ spl4_18 | spl4_52 | ~ spl4_80)), inference(subsumption_resolution, [], [f5402, f827])).
fof(f827, plain, (~ (e1 = op(e2, e4)) | spl4_52), inference(avatar_component_clause, [], [f826])).
fof(f5402, plain, ((e1 = op(e2, e4)) | (~ spl4_18 | ~ spl4_80)), inference(backward_demodulation, [], [f5029, f685])).
fof(f5382, plain, (~ spl4_31 | spl4_124), inference(avatar_contradiction_clause, [], [f5381])).
fof(f5381, plain, ($false | (~ spl4_31 | spl4_124)), inference(subsumption_resolution, [], [f5368, f1129])).
fof(f1129, plain, (~ (op(e0, e0) = e3) | spl4_124), inference(avatar_component_clause, [], [f1128])).
fof(f5368, plain, ((op(e0, e0) = e3) | ~ spl4_31), inference(backward_demodulation, [], [f349, f740])).
fof(f5367, plain, (~ spl4_40 | ~ spl4_57 | spl4_79), inference(avatar_contradiction_clause, [], [f5366])).
fof(f5366, plain, ($false | (~ spl4_40 | ~ spl4_57 | spl4_79)), inference(subsumption_resolution, [], [f5364, f940])).
fof(f940, plain, (~ (e3 = op(e1, e4)) | spl4_79), inference(avatar_component_clause, [], [f939])).
fof(f5364, plain, ((e3 = op(e1, e4)) | (~ spl4_40 | ~ spl4_57)), inference(backward_demodulation, [], [f5322, f777])).
fof(f5365, plain, (spl4_18 | ~ spl4_40 | ~ spl4_57), inference(avatar_split_clause, [], [f5363, f847, f775, f683])).
fof(f5363, plain, ((e2 = op(e4, e1)) | (~ spl4_40 | ~ spl4_57)), inference(backward_demodulation, [], [f5321, f777])).
fof(f5356, plain, (~ spl4_18 | ~ spl4_43), inference(avatar_split_clause, [], [f5348, f788, f683])).
fof(f5348, plain, (~ (e2 = op(e4, e1)) | ~ spl4_43), inference(backward_demodulation, [], [f375, f790])).
fof(f5329, plain, (~ spl4_32 | ~ spl4_57), inference(avatar_split_clause, [], [f5323, f847, f742])).
fof(f5323, plain, (~ (e1 = op(e3, e3)) | ~ spl4_57), inference(backward_demodulation, [], [f393, f849])).
fof(f5296, plain, (~ spl4_69 | ~ spl4_70), inference(avatar_contradiction_clause, [], [f5295])).
fof(f5295, plain, ($false | (~ spl4_69 | ~ spl4_70)), inference(subsumption_resolution, [], [f5294, f465])).
fof(f5294, plain, ((e3 = e4) | (~ spl4_69 | ~ spl4_70)), inference(backward_demodulation, [], [f903, f899])).
fof(f5293, plain, (~ spl4_70 | ~ spl4_86 | spl4_103), inference(avatar_contradiction_clause, [], [f5292])).
fof(f5292, plain, ($false | (~ spl4_70 | ~ spl4_86 | spl4_103)), inference(subsumption_resolution, [], [f5289, f1041])).
fof(f1041, plain, (~ (e2 = op(e0, e4)) | spl4_103), inference(avatar_component_clause, [], [f1040])).
fof(f5289, plain, ((e2 = op(e0, e4)) | (~ spl4_70 | ~ spl4_86)), inference(backward_demodulation, [], [f5021, f903])).
fof(f5021, plain, ((e2 = op(e0, op(e2, e1))) | ~ spl4_86), inference(backward_demodulation, [], [f342, f971])).
fof(f5282, plain, (~ spl4_51 | ~ spl4_71), inference(avatar_split_clause, [], [f5274, f906, f822])).
fof(f5274, plain, (~ (e0 = op(e2, e4)) | ~ spl4_71), inference(backward_demodulation, [], [f429, f908])).
fof(f5269, plain, (~ spl4_9 | ~ spl4_84), inference(avatar_split_clause, [], [f5267, f960, f645])).
fof(f5267, plain, (~ (e3 = op(e4, e3)) | ~ spl4_84), inference(backward_demodulation, [], [f392, f962])).
fof(f5238, plain, (~ spl4_32 | ~ spl4_104 | ~ spl4_125 | spl4_164), inference(avatar_split_clause, [], [f5229, f1396, f1132, f1044, f742])).
fof(f5229, plain, (~ (e1 = op(e3, e3)) | (~ spl4_104 | ~ spl4_125 | spl4_164)), inference(backward_demodulation, [], [f4947, f1046])).
fof(f4947, plain, (~ (e1 = op(op(e0, e4), op(e0, e4))) | (~ spl4_125 | spl4_164)), inference(backward_demodulation, [], [f1398, f1134])).
fof(f5237, plain, (spl4_45 | ~ spl4_22 | ~ spl4_104), inference(avatar_split_clause, [], [f5228, f1044, f700, f796])).
fof(f5228, plain, ((e4 = op(e3, e1)) | (~ spl4_22 | ~ spl4_104)), inference(backward_demodulation, [], [f4904, f1046])).
fof(f4904, plain, ((e4 = op(op(e0, e4), e1)) | ~ spl4_22), inference(forward_demodulation, [], [f351, f702])).
fof(f5236, plain, (~ spl4_22 | spl4_81 | ~ spl4_104), inference(avatar_contradiction_clause, [], [f5235])).
fof(f5235, plain, ($false | (~ spl4_22 | spl4_81 | ~ spl4_104)), inference(subsumption_resolution, [], [f5227, f949])).
fof(f5227, plain, ((e0 = op(e1, e3)) | (~ spl4_22 | ~ spl4_104)), inference(backward_demodulation, [], [f4903, f1046])).
fof(f5234, plain, (~ spl4_29 | ~ spl4_104), inference(avatar_split_clause, [], [f5225, f1044, f729])).
fof(f5225, plain, (~ (e3 = op(e3, e4)) | ~ spl4_104), inference(backward_demodulation, [], [f398, f1046])).
fof(f5223, plain, (~ spl4_58 | ~ spl4_108), inference(avatar_split_clause, [], [f5217, f1061, f851])).
fof(f5217, plain, (~ (e2 = op(e2, e3)) | ~ spl4_108), inference(backward_demodulation, [], [f387, f1063])).
fof(f5172, plain, (~ spl4_104 | ~ spl4_125 | spl4_168), inference(avatar_split_clause, [], [f4948, f1414, f1132, f1044])).
fof(f4948, plain, (~ (e3 = op(e0, e4)) | (~ spl4_125 | spl4_168)), inference(backward_demodulation, [], [f1416, f1134])).
fof(f5168, plain, (~ spl4_85 | ~ spl4_80), inference(avatar_split_clause, [], [f5167, f943, f964])).
fof(f5167, plain, (~ (e4 = op(e1, e3)) | ~ spl4_80), inference(forward_demodulation, [], [f425, f945])).
fof(f5160, plain, (spl4_65 | ~ spl4_3), inference(avatar_split_clause, [], [f4609, f620, f880])).
fof(f4609, plain, ((e4 = op(e2, e2)) | ~ spl4_3), inference(backward_demodulation, [], [f355, f622])).
fof(f5158, plain, (~ spl4_18 | ~ spl4_3), inference(avatar_split_clause, [], [f4902, f620, f683])).
fof(f4902, plain, (~ (e2 = op(e4, e1)) | ~ spl4_3), inference(forward_demodulation, [], [f452, f622])).
fof(f5156, plain, (~ spl4_9 | ~ spl4_29 | ~ spl4_32), inference(avatar_contradiction_clause, [], [f5155])).
fof(f5155, plain, ($false | (~ spl4_9 | ~ spl4_29 | ~ spl4_32)), inference(subsumption_resolution, [], [f5154, f462])).
fof(f5154, plain, ((e1 = e4) | (~ spl4_9 | ~ spl4_29 | ~ spl4_32)), inference(forward_demodulation, [], [f5150, f744])).
fof(f5150, plain, ((e4 = op(e3, e3)) | (~ spl4_9 | ~ spl4_29)), inference(backward_demodulation, [], [f4880, f647])).
fof(f4880, plain, ((e4 = op(e3, op(e4, e3))) | ~ spl4_29), inference(backward_demodulation, [], [f354, f731])).
fof(f5147, plain, (~ spl4_12 | ~ spl4_51 | ~ spl4_118), inference(avatar_contradiction_clause, [], [f5146])).
fof(f5146, plain, ($false | (~ spl4_12 | ~ spl4_51 | ~ spl4_118)), inference(subsumption_resolution, [], [f5145, f464])).
fof(f5145, plain, ((e2 = e4) | (~ spl4_12 | ~ spl4_51 | ~ spl4_118)), inference(forward_demodulation, [], [f5139, f1105])).
fof(f5139, plain, ((e4 = op(e0, e1)) | (~ spl4_12 | ~ spl4_51)), inference(backward_demodulation, [], [f5063, f660])).
fof(f5063, plain, ((e4 = op(e0, op(e4, e2))) | ~ spl4_51), inference(backward_demodulation, [], [f353, f824])).
fof(f5144, plain, (~ spl4_3 | ~ spl4_12 | ~ spl4_94 | spl4_183), inference(avatar_contradiction_clause, [], [f5143])).
fof(f5143, plain, ($false | (~ spl4_3 | ~ spl4_12 | ~ spl4_94 | spl4_183)), inference(subsumption_resolution, [], [f5136, f1004])).
fof(f5136, plain, (~ (e3 = op(e1, e1)) | (~ spl4_3 | ~ spl4_12 | spl4_183)), inference(backward_demodulation, [], [f4688, f660])).
fof(f4688, plain, (~ (e3 = op(op(e4, e2), op(e4, e2))) | (~ spl4_3 | spl4_183)), inference(forward_demodulation, [], [f1515, f622])).
fof(f5142, plain, (~ spl4_3 | ~ spl4_12 | ~ spl4_16 | spl4_161), inference(avatar_contradiction_clause, [], [f5141])).
fof(f5141, plain, ($false | (~ spl4_3 | ~ spl4_12 | ~ spl4_16 | spl4_161)), inference(subsumption_resolution, [], [f5135, f677])).
fof(f5135, plain, (~ (e0 = op(e4, e1)) | (~ spl4_3 | ~ spl4_12 | spl4_161)), inference(backward_demodulation, [], [f4615, f660])).
fof(f4615, plain, (~ (e0 = op(e4, op(e4, e2))) | (~ spl4_3 | spl4_161)), inference(backward_demodulation, [], [f1379, f622])).
fof(f5117, plain, (~ spl4_22 | ~ spl4_25), inference(avatar_contradiction_clause, [], [f5116])).
fof(f5116, plain, ($false | (~ spl4_22 | ~ spl4_25)), inference(subsumption_resolution, [], [f5115, f462])).
fof(f5115, plain, ((e1 = e4) | (~ spl4_22 | ~ spl4_25)), inference(forward_demodulation, [], [f714, f702])).
fof(f5094, plain, (spl4_2 | ~ spl4_45 | ~ spl4_85), inference(avatar_contradiction_clause, [], [f5093])).
fof(f5093, plain, ($false | (spl4_2 | ~ spl4_45 | ~ spl4_85)), inference(subsumption_resolution, [], [f5090, f617])).
fof(f617, plain, (~ (e1 = op(e4, e4)) | spl4_2), inference(avatar_component_clause, [], [f616])).
fof(f5090, plain, ((e1 = op(e4, e4)) | (~ spl4_45 | ~ spl4_85)), inference(backward_demodulation, [], [f4259, f798])).
fof(f4259, plain, ((e1 = op(op(e3, e1), e4)) | ~ spl4_85), inference(backward_demodulation, [], [f339, f966])).
fof(f5092, plain, (~ spl4_20 | ~ spl4_45), inference(avatar_split_clause, [], [f5088, f796, f691])).
fof(f5088, plain, (~ (e4 = op(e4, e1)) | ~ spl4_45), inference(backward_demodulation, [], [f375, f798])).
fof(f5085, plain, (~ spl4_46 | ~ spl4_106 | spl4_121), inference(avatar_contradiction_clause, [], [f5084])).
fof(f5084, plain, ($false | (~ spl4_46 | ~ spl4_106 | spl4_121)), inference(subsumption_resolution, [], [f5081, f1117])).
fof(f1117, plain, (~ (e0 = op(e0, e0)) | spl4_121), inference(avatar_component_clause, [], [f1116])).
fof(f5081, plain, ((e0 = op(e0, e0)) | (~ spl4_46 | ~ spl4_106)), inference(backward_demodulation, [], [f4984, f803])).
fof(f4984, plain, ((e0 = op(op(e3, e0), e0)) | ~ spl4_106), inference(backward_demodulation, [], [f334, f1055])).
fof(f5048, plain, (~ spl4_67 | ~ spl4_86 | spl4_97), inference(avatar_contradiction_clause, [], [f5047])).
fof(f5047, plain, ($false | (~ spl4_67 | ~ spl4_86 | spl4_97)), inference(subsumption_resolution, [], [f5044, f1016])).
fof(f1016, plain, (~ (e1 = op(e1, e0)) | spl4_97), inference(avatar_component_clause, [], [f1015])).
fof(f5044, plain, ((e1 = op(e1, e0)) | (~ spl4_67 | ~ spl4_86)), inference(backward_demodulation, [], [f5020, f891])).
fof(f5020, plain, ((e1 = op(op(e2, e1), e0)) | ~ spl4_86), inference(backward_demodulation, [], [f338, f971])).
fof(f5036, plain, (~ spl4_77 | ~ spl4_80), inference(avatar_contradiction_clause, [], [f5035])).
fof(f5035, plain, ($false | (~ spl4_77 | ~ spl4_80)), inference(subsumption_resolution, [], [f5034, f462])).
fof(f5034, plain, ((e1 = e4) | (~ spl4_77 | ~ spl4_80)), inference(backward_demodulation, [], [f945, f933])).
fof(f5007, plain, (~ spl4_22 | spl4_95 | ~ spl4_102), inference(avatar_contradiction_clause, [], [f5006])).
fof(f5006, plain, ($false | (~ spl4_22 | spl4_95 | ~ spl4_102)), inference(subsumption_resolution, [], [f5000, f1007])).
fof(f1007, plain, (~ (e4 = op(e1, e1)) | spl4_95), inference(avatar_component_clause, [], [f1006])).
fof(f5000, plain, ((e4 = op(e1, e1)) | (~ spl4_22 | ~ spl4_102)), inference(backward_demodulation, [], [f4904, f1038])).
fof(f5005, plain, (~ spl4_22 | spl4_91 | ~ spl4_102), inference(avatar_contradiction_clause, [], [f5004])).
fof(f5004, plain, ($false | (~ spl4_22 | spl4_91 | ~ spl4_102)), inference(subsumption_resolution, [], [f4999, f991])).
fof(f4999, plain, ((e0 = op(e1, e1)) | (~ spl4_22 | ~ spl4_102)), inference(backward_demodulation, [], [f4903, f1038])).
fof(f4980, plain, (~ spl4_14 | ~ spl4_114), inference(avatar_split_clause, [], [f4974, f1086, f666])).
fof(f4974, plain, (~ (e3 = op(e4, e2)) | ~ spl4_114), inference(backward_demodulation, [], [f379, f1088])).
fof(f4972, plain, (spl4_61 | ~ spl4_98 | ~ spl4_118), inference(avatar_contradiction_clause, [], [f4971])).
fof(f4971, plain, ($false | (spl4_61 | ~ spl4_98 | ~ spl4_118)), inference(subsumption_resolution, [], [f4966, f865])).
fof(f865, plain, (~ (e0 = op(e2, e2)) | spl4_61), inference(avatar_component_clause, [], [f864])).
fof(f4966, plain, ((e0 = op(e2, e2)) | (~ spl4_98 | ~ spl4_118)), inference(backward_demodulation, [], [f4685, f1105])).
fof(f4685, plain, ((e0 = op(e2, op(e0, e1))) | ~ spl4_98), inference(forward_demodulation, [], [f332, f1021])).
fof(f4970, plain, (spl4_62 | ~ spl4_98 | ~ spl4_118), inference(avatar_contradiction_clause, [], [f4969])).
fof(f4969, plain, ($false | (spl4_62 | ~ spl4_98 | ~ spl4_118)), inference(subsumption_resolution, [], [f4965, f869])).
fof(f4965, plain, ((e1 = op(e2, e2)) | (~ spl4_98 | ~ spl4_118)), inference(backward_demodulation, [], [f4684, f1105])).
fof(f4684, plain, ((e1 = op(op(e0, e1), e2)) | ~ spl4_98), inference(forward_demodulation, [], [f336, f1021])).
fof(f4967, plain, (~ spl4_43 | ~ spl4_118), inference(avatar_split_clause, [], [f4958, f1103, f788])).
fof(f4958, plain, (~ (e2 = op(e3, e1)) | ~ spl4_118), inference(backward_demodulation, [], [f368, f1105])).
fof(f4955, plain, (~ spl4_102 | ~ spl4_125 | spl4_180), inference(avatar_split_clause, [], [f4950, f1492, f1132, f1036])).
fof(f4950, plain, (~ (e1 = op(e0, e4)) | (~ spl4_125 | spl4_180)), inference(backward_demodulation, [], [f1494, f1134])).
fof(f1494, plain, (~ (e1 = op(e0, op(e0, e0))) | spl4_180), inference(avatar_component_clause, [], [f1492])).
fof(f4936, plain, (~ spl4_76 | ~ spl4_146 | spl4_177), inference(avatar_split_clause, [], [f4935, f1466, f1300, f927])).
fof(f4935, plain, (~ (e0 = op(e1, e4)) | (~ spl4_146 | spl4_177)), inference(forward_demodulation, [], [f1468, f1301])).
fof(f4888, plain, (~ spl4_21 | ~ spl4_22), inference(avatar_contradiction_clause, [], [f4887])).
fof(f4887, plain, ($false | (~ spl4_21 | ~ spl4_22)), inference(subsumption_resolution, [], [f4886, f456])).
fof(f4886, plain, ((e0 = e1) | (~ spl4_21 | ~ spl4_22)), inference(forward_demodulation, [], [f702, f698])).
fof(f4885, plain, (~ spl4_27 | ~ spl4_29), inference(avatar_contradiction_clause, [], [f4884])).
fof(f4884, plain, ($false | (~ spl4_27 | ~ spl4_29)), inference(subsumption_resolution, [], [f4883, f461])).
fof(f4883, plain, ((e1 = e3) | (~ spl4_27 | ~ spl4_29)), inference(backward_demodulation, [], [f731, f723])).
fof(f4876, plain, (~ spl4_3 | ~ spl4_31 | ~ spl4_50 | spl4_174), inference(avatar_contradiction_clause, [], [f4875])).
fof(f4875, plain, ($false | (~ spl4_3 | ~ spl4_31 | ~ spl4_50 | spl4_174)), inference(subsumption_resolution, [], [f4874, f622])).
fof(f4874, plain, (~ (e2 = op(e4, e4)) | (~ spl4_31 | ~ spl4_50 | spl4_174)), inference(forward_demodulation, [], [f4867, f819])).
fof(f4867, plain, (~ (e2 = op(op(e3, e0), op(e3, e0))) | (~ spl4_31 | spl4_174)), inference(backward_demodulation, [], [f1454, f740])).
fof(f4871, plain, (~ spl4_31 | ~ spl4_123), inference(avatar_contradiction_clause, [], [f4870])).
fof(f4870, plain, ($false | (~ spl4_31 | ~ spl4_123)), inference(subsumption_resolution, [], [f4869, f463])).
fof(f4869, plain, ((e2 = e3) | (~ spl4_31 | ~ spl4_123)), inference(forward_demodulation, [], [f4857, f1126])).
fof(f4857, plain, ((op(e0, e0) = e3) | ~ spl4_31), inference(backward_demodulation, [], [f349, f740])).
fof(f4799, plain, (~ spl4_84 | ~ spl4_85), inference(avatar_contradiction_clause, [], [f4798])).
fof(f4798, plain, ($false | (~ spl4_84 | ~ spl4_85)), inference(subsumption_resolution, [], [f4797, f465])).
fof(f4797, plain, ((e3 = e4) | (~ spl4_84 | ~ spl4_85)), inference(backward_demodulation, [], [f966, f962])).
fof(f4787, plain, (~ spl4_92 | ~ spl4_146), inference(avatar_contradiction_clause, [], [f4786])).
fof(f4786, plain, ($false | (~ spl4_92 | ~ spl4_146)), inference(subsumption_resolution, [], [f4785, f462])).
fof(f4785, plain, ((e1 = e4) | (~ spl4_92 | ~ spl4_146)), inference(forward_demodulation, [], [f4781, f996])).
fof(f4781, plain, ((e4 = op(e1, e1)) | (~ spl4_92 | ~ spl4_146)), inference(backward_demodulation, [], [f1301, f996])).
fof(f4771, plain, (~ spl4_21 | ~ spl4_105), inference(avatar_contradiction_clause, [], [f4770])).
fof(f4770, plain, ($false | (~ spl4_21 | ~ spl4_105)), inference(subsumption_resolution, [], [f4769, f459])).
fof(f4769, plain, ((e0 = e4) | (~ spl4_21 | ~ spl4_105)), inference(forward_demodulation, [], [f4766, f698])).
fof(f4766, plain, ((e4 = op(e4, e0)) | (~ spl4_21 | ~ spl4_105)), inference(backward_demodulation, [], [f4674, f1050])).
fof(f4674, plain, ((e4 = op(op(e0, e4), e0)) | ~ spl4_21), inference(forward_demodulation, [], [f351, f698])).
fof(f4768, plain, (~ spl4_30 | ~ spl4_105), inference(avatar_split_clause, [], [f4765, f1048, f733])).
fof(f4765, plain, (~ (e4 = op(e3, e4)) | ~ spl4_105), inference(backward_demodulation, [], [f398, f1050])).
fof(f4762, plain, (~ spl4_50 | spl4_54 | ~ spl4_108), inference(avatar_contradiction_clause, [], [f4761])).
fof(f4761, plain, ($false | (~ spl4_50 | spl4_54 | ~ spl4_108)), inference(subsumption_resolution, [], [f4756, f835])).
fof(f835, plain, (~ (e3 = op(e2, e4)) | spl4_54), inference(avatar_component_clause, [], [f834])).
fof(f4756, plain, ((e3 = op(e2, e4)) | (~ spl4_50 | ~ spl4_108)), inference(backward_demodulation, [], [f3172, f1063])).
fof(f3172, plain, ((e3 = op(op(e0, e3), e4)) | ~ spl4_50), inference(forward_demodulation, [], [f346, f819])).
fof(f4758, plain, (~ spl4_8 | ~ spl4_108), inference(avatar_split_clause, [], [f4751, f1061, f641])).
fof(f4751, plain, (~ (e2 = op(e4, e3)) | ~ spl4_108), inference(backward_demodulation, [], [f389, f1063])).
fof(f4740, plain, (~ spl4_74 | spl4_83 | ~ spl4_112), inference(avatar_contradiction_clause, [], [f4739])).
fof(f4739, plain, ($false | (~ spl4_74 | spl4_83 | ~ spl4_112)), inference(subsumption_resolution, [], [f4733, f957])).
fof(f4733, plain, ((e2 = op(e1, e3)) | (~ spl4_74 | ~ spl4_112)), inference(backward_demodulation, [], [f4677, f1080])).
fof(f4677, plain, ((e2 = op(op(e0, e2), e3)) | ~ spl4_74), inference(forward_demodulation, [], [f341, f920])).
fof(f4738, plain, (~ spl4_37 | ~ spl4_112), inference(avatar_split_clause, [], [f4727, f1078, f763])).
fof(f4727, plain, (~ (e1 = op(e3, e2)) | ~ spl4_112), inference(backward_demodulation, [], [f378, f1080])).
fof(f4721, plain, (~ spl4_94 | ~ spl4_119), inference(avatar_split_clause, [], [f4713, f1107, f1002])).
fof(f4713, plain, (~ (e3 = op(e1, e1)) | ~ spl4_119), inference(backward_demodulation, [], [f366, f1109])).
fof(f4687, plain, (~ spl4_123 | ~ spl4_98), inference(avatar_split_clause, [], [f4501, f1019, f1124])).
fof(f4501, plain, (~ (op(e0, e0) = e2) | ~ spl4_98), inference(forward_demodulation, [], [f356, f1021])).
fof(f4669, plain, (~ spl4_3 | ~ spl4_146 | spl4_176), inference(avatar_contradiction_clause, [], [f4668])).
fof(f4668, plain, ($false | (~ spl4_3 | ~ spl4_146 | spl4_176)), inference(subsumption_resolution, [], [f4667, f622])).
fof(f4667, plain, (~ (e2 = op(e4, e4)) | (~ spl4_146 | spl4_176)), inference(forward_demodulation, [], [f1464, f1301])).
fof(f4652, plain, (~ spl4_78 | ~ spl4_98), inference(avatar_split_clause, [], [f4651, f1019, f935])).
fof(f4651, plain, (~ (e2 = op(e1, e4)) | ~ spl4_98), inference(forward_demodulation, [], [f419, f1021])).
fof(f4644, plain, (~ spl4_54 | ~ spl4_74), inference(avatar_split_clause, [], [f4643, f918, f834])).
fof(f4643, plain, (~ (e3 = op(e2, e4)) | ~ spl4_74), inference(forward_demodulation, [], [f429, f920])).
fof(f4642, plain, (spl4_52 | ~ spl4_43 | ~ spl4_85), inference(avatar_split_clause, [], [f4571, f964, f788, f826])).
fof(f4571, plain, ((e1 = op(e2, e4)) | (~ spl4_43 | ~ spl4_85)), inference(backward_demodulation, [], [f4259, f790])).
fof(f4637, plain, (~ spl4_35 | ~ spl4_50), inference(avatar_split_clause, [], [f4636, f817, f754])).
fof(f4636, plain, (~ (e4 = op(e3, e3)) | ~ spl4_50), inference(forward_demodulation, [], [f438, f819])).
fof(f4635, plain, (~ spl4_35 | ~ spl4_85), inference(avatar_split_clause, [], [f4634, f964, f754])).
fof(f4634, plain, (~ (e4 = op(e3, e3)) | ~ spl4_85), inference(forward_demodulation, [], [f391, f966])).
fof(f4633, plain, (~ spl4_31 | ~ spl4_3 | ~ spl4_14 | spl4_133), inference(avatar_split_clause, [], [f4618, f1241, f666, f620, f738])).
fof(f4618, plain, (~ (e0 = op(e3, e3)) | (~ spl4_3 | ~ spl4_14 | spl4_133)), inference(forward_demodulation, [], [f4612, f668])).
fof(f4612, plain, (~ (e0 = op(op(e4, e2), op(e4, e2))) | (~ spl4_3 | spl4_133)), inference(backward_demodulation, [], [f1243, f622])).
fof(f4620, plain, (~ spl4_3 | ~ spl4_14 | spl4_135), inference(avatar_contradiction_clause, [], [f4619])).
fof(f4619, plain, ($false | (~ spl4_3 | ~ spl4_14 | spl4_135)), inference(subsumption_resolution, [], [f4613, f668])).
fof(f4613, plain, (~ (e3 = op(e4, e2)) | (~ spl4_3 | spl4_135)), inference(backward_demodulation, [], [f1251, f622])).
fof(f4608, plain, (~ spl4_7 | ~ spl4_30 | ~ spl4_77), inference(avatar_contradiction_clause, [], [f4607])).
fof(f4607, plain, ($false | (~ spl4_7 | ~ spl4_30 | ~ spl4_77)), inference(subsumption_resolution, [], [f4606, f461])).
fof(f4606, plain, ((e1 = e3) | (~ spl4_7 | ~ spl4_30 | ~ spl4_77)), inference(forward_demodulation, [], [f4605, f933])).
fof(f4605, plain, ((e3 = op(e1, e4)) | (~ spl4_7 | ~ spl4_30)), inference(backward_demodulation, [], [f4435, f639])).
fof(f4435, plain, ((e3 = op(op(e4, e3), e4)) | ~ spl4_30), inference(forward_demodulation, [], [f350, f735])).
fof(f4589, plain, (~ spl4_20 | ~ spl4_77), inference(avatar_contradiction_clause, [], [f4588])).
fof(f4588, plain, ($false | (~ spl4_20 | ~ spl4_77)), inference(subsumption_resolution, [], [f4587, f462])).
fof(f4587, plain, ((e1 = e4) | (~ spl4_20 | ~ spl4_77)), inference(forward_demodulation, [], [f4584, f933])).
fof(f4584, plain, ((e4 = op(e1, e4)) | (~ spl4_20 | ~ spl4_77)), inference(backward_demodulation, [], [f4417, f693])).
fof(f4417, plain, ((e4 = op(e1, op(e4, e1))) | ~ spl4_77), inference(forward_demodulation, [], [f352, f933])).
fof(f4579, plain, (~ spl4_36 | ~ spl4_37), inference(avatar_contradiction_clause, [], [f4578])).
fof(f4578, plain, ($false | (~ spl4_36 | ~ spl4_37)), inference(subsumption_resolution, [], [f4577, f456])).
fof(f4577, plain, ((e0 = e1) | (~ spl4_36 | ~ spl4_37)), inference(backward_demodulation, [], [f765, f761])).
fof(f4576, plain, (spl4_14 | ~ spl4_43 | ~ spl4_85), inference(avatar_split_clause, [], [f4572, f964, f788, f666])).
fof(f4572, plain, ((e3 = op(e4, e2)) | (~ spl4_43 | ~ spl4_85)), inference(backward_demodulation, [], [f4260, f790])).
fof(f4260, plain, ((e3 = op(e4, op(e3, e1))) | ~ spl4_85), inference(backward_demodulation, [], [f347, f966])).
fof(f4575, plain, (~ spl4_43 | ~ spl4_51 | ~ spl4_85), inference(avatar_contradiction_clause, [], [f4574])).
fof(f4574, plain, ($false | (~ spl4_43 | ~ spl4_51 | ~ spl4_85)), inference(subsumption_resolution, [], [f4573, f456])).
fof(f4573, plain, ((e0 = e1) | (~ spl4_43 | ~ spl4_51 | ~ spl4_85)), inference(forward_demodulation, [], [f4571, f824])).
fof(f4567, plain, (~ spl4_37 | ~ spl4_58 | spl4_88), inference(avatar_contradiction_clause, [], [f4566])).
fof(f4566, plain, ($false | (~ spl4_37 | ~ spl4_58 | spl4_88)), inference(subsumption_resolution, [], [f4561, f978])).
fof(f978, plain, (~ (e2 = op(e1, e2)) | spl4_88), inference(avatar_component_clause, [], [f977])).
fof(f4561, plain, ((e2 = op(e1, e2)) | (~ spl4_37 | ~ spl4_58)), inference(backward_demodulation, [], [f4464, f853])).
fof(f4464, plain, ((e2 = op(e1, op(e2, e3))) | ~ spl4_37), inference(forward_demodulation, [], [f344, f765])).
fof(f4565, plain, (~ spl4_37 | ~ spl4_58 | ~ spl4_67), inference(avatar_contradiction_clause, [], [f4564])).
fof(f4564, plain, ($false | (~ spl4_37 | ~ spl4_58 | ~ spl4_67)), inference(subsumption_resolution, [], [f4563, f461])).
fof(f4563, plain, ((e1 = e3) | (~ spl4_37 | ~ spl4_58 | ~ spl4_67)), inference(forward_demodulation, [], [f4560, f891])).
fof(f4560, plain, ((e3 = op(e2, e1)) | (~ spl4_37 | ~ spl4_58)), inference(backward_demodulation, [], [f4463, f853])).
fof(f4463, plain, ((e3 = op(op(e2, e3), e1)) | ~ spl4_37), inference(forward_demodulation, [], [f348, f765])).
fof(f4556, plain, (~ spl4_67 | spl4_82 | ~ spl4_89), inference(avatar_contradiction_clause, [], [f4555])).
fof(f4555, plain, ($false | (~ spl4_67 | spl4_82 | ~ spl4_89)), inference(subsumption_resolution, [], [f4552, f953])).
fof(f953, plain, (~ (e1 = op(e1, e3)) | spl4_82), inference(avatar_component_clause, [], [f952])).
fof(f4552, plain, ((e1 = op(e1, e3)) | (~ spl4_67 | ~ spl4_89)), inference(backward_demodulation, [], [f4535, f891])).
fof(f4535, plain, ((e1 = op(op(e2, e1), e3)) | ~ spl4_89), inference(backward_demodulation, [], [f338, f983])).
fof(f4531, plain, (~ spl4_91 | ~ spl4_94), inference(avatar_contradiction_clause, [], [f4530])).
fof(f4530, plain, ($false | (~ spl4_91 | ~ spl4_94)), inference(subsumption_resolution, [], [f4529, f458])).
fof(f4529, plain, ((e0 = e3) | (~ spl4_91 | ~ spl4_94)), inference(backward_demodulation, [], [f1004, f992])).
fof(f4522, plain, (~ spl4_104 | ~ spl4_120 | ~ spl4_122 | spl4_170), inference(avatar_contradiction_clause, [], [f4521])).
fof(f4521, plain, ($false | (~ spl4_104 | ~ spl4_120 | ~ spl4_122 | spl4_170)), inference(subsumption_resolution, [], [f4520, f1046])).
fof(f4520, plain, (~ (e3 = op(e0, e4)) | (~ spl4_120 | ~ spl4_122 | spl4_170)), inference(forward_demodulation, [], [f4515, f1113])).
fof(f4515, plain, (~ (e3 = op(e0, op(e0, e1))) | (~ spl4_122 | spl4_170)), inference(backward_demodulation, [], [f1428, f1122])).
fof(f4499, plain, (~ spl4_125 | ~ spl4_120), inference(avatar_split_clause, [], [f4498, f1111, f1132])).
fof(f4498, plain, (~ (op(e0, e0) = e4) | ~ spl4_120), inference(forward_demodulation, [], [f406, f1113])).
fof(f4481, plain, (spl4_12 | ~ spl4_98 | ~ spl4_120), inference(avatar_split_clause, [], [f4480, f1111, f1019, f658])).
fof(f4480, plain, ((e1 = op(e4, e2)) | (~ spl4_98 | ~ spl4_120)), inference(forward_demodulation, [], [f4172, f1021])).
fof(f4172, plain, ((e1 = op(e4, op(e1, e0))) | ~ spl4_120), inference(backward_demodulation, [], [f336, f1113])).
fof(f4479, plain, (~ spl4_89 | ~ spl4_94), inference(avatar_split_clause, [], [f4234, f1002, f981])).
fof(f4234, plain, (~ (e3 = op(e1, e2)) | ~ spl4_94), inference(backward_demodulation, [], [f420, f1004])).
fof(f4460, plain, (spl4_32 | ~ spl4_34 | ~ spl4_158), inference(avatar_split_clause, [], [f4459, f1364, f750, f742])).
fof(f4459, plain, ((e1 = op(e3, e3)) | (~ spl4_34 | ~ spl4_158)), inference(forward_demodulation, [], [f4458, f752])).
fof(f4458, plain, ((e1 = op(op(e3, e3), op(e3, e3))) | (~ spl4_34 | ~ spl4_158)), inference(forward_demodulation, [], [f1365, f752])).
fof(f1365, plain, ((e1 = op(op(e3, op(e3, e3)), op(e3, op(e3, e3)))) | ~ spl4_158), inference(avatar_component_clause, [], [f1364])).
fof(f4430, plain, (~ spl4_35 | ~ spl4_30), inference(avatar_split_clause, [], [f4429, f733, f754])).
fof(f4429, plain, (~ (e4 = op(e3, e3)) | ~ spl4_30), inference(forward_demodulation, [], [f445, f735])).
fof(f4407, plain, (~ spl4_2 | ~ spl4_77), inference(avatar_split_clause, [], [f4406, f931, f616])).
fof(f4406, plain, (~ (e1 = op(e4, e4)) | ~ spl4_77), inference(forward_demodulation, [], [f402, f933])).
fof(f4405, plain, (~ spl4_3 | ~ spl4_85 | ~ spl4_94 | spl4_176), inference(avatar_split_clause, [], [f4265, f1462, f1002, f964, f620])).
fof(f4265, plain, (~ (e2 = op(e4, e4)) | (~ spl4_85 | ~ spl4_94 | spl4_176)), inference(backward_demodulation, [], [f4240, f966])).
fof(f4240, plain, (~ (e2 = op(op(e1, e3), op(e1, e3))) | (~ spl4_94 | spl4_176)), inference(backward_demodulation, [], [f1464, f1004])).
fof(f4401, plain, (~ spl4_8 | ~ spl4_27 | ~ spl4_86), inference(avatar_contradiction_clause, [], [f4400])).
fof(f4400, plain, ($false | (~ spl4_8 | ~ spl4_27 | ~ spl4_86)), inference(subsumption_resolution, [], [f4399, f459])).
fof(f4399, plain, ((e0 = e4) | (~ spl4_8 | ~ spl4_27 | ~ spl4_86)), inference(forward_demodulation, [], [f4395, f971])).
fof(f4395, plain, ((e4 = op(e1, e2)) | (~ spl4_8 | ~ spl4_27)), inference(backward_demodulation, [], [f4011, f643])).
fof(f4011, plain, ((e4 = op(e1, op(e4, e3))) | ~ spl4_27), inference(backward_demodulation, [], [f354, f723])).
fof(f4388, plain, (~ spl4_15 | ~ spl4_21 | ~ spl4_51), inference(avatar_contradiction_clause, [], [f4387])).
fof(f4387, plain, ($false | (~ spl4_15 | ~ spl4_21 | ~ spl4_51)), inference(subsumption_resolution, [], [f4386, f457])).
fof(f4386, plain, ((e0 = e2) | (~ spl4_15 | ~ spl4_21 | ~ spl4_51)), inference(forward_demodulation, [], [f4384, f698])).
fof(f4384, plain, ((e2 = op(e4, e0)) | (~ spl4_15 | ~ spl4_51)), inference(backward_demodulation, [], [f4314, f672])).
fof(f4314, plain, ((e2 = op(op(e4, e2), e0)) | ~ spl4_51), inference(backward_demodulation, [], [f345, f824])).
fof(f4382, plain, (~ spl4_17 | ~ spl4_19), inference(avatar_contradiction_clause, [], [f4381])).
fof(f4381, plain, ($false | (~ spl4_17 | ~ spl4_19)), inference(subsumption_resolution, [], [f4380, f461])).
fof(f4380, plain, ((e1 = e3) | (~ spl4_17 | ~ spl4_19)), inference(forward_demodulation, [], [f689, f681])).
fof(f4373, plain, (~ spl4_6 | ~ spl4_21), inference(avatar_split_clause, [], [f4362, f696, f633])).
fof(f4362, plain, (~ (e0 = op(e4, e3)) | ~ spl4_21), inference(backward_demodulation, [], [f448, f698])).
fof(f4360, plain, (~ spl4_27 | ~ spl4_28), inference(avatar_contradiction_clause, [], [f4359])).
fof(f4359, plain, ($false | (~ spl4_27 | ~ spl4_28)), inference(subsumption_resolution, [], [f4358, f460])).
fof(f460, plain, ~ (e1 = e2), inference(cnf_transformation, [], [f5])).
fof(f4358, plain, ((e1 = e2) | (~ spl4_27 | ~ spl4_28)), inference(forward_demodulation, [], [f727, f723])).
fof(f4357, plain, (~ spl4_27 | ~ spl4_30), inference(avatar_contradiction_clause, [], [f4356])).
fof(f4356, plain, ($false | (~ spl4_27 | ~ spl4_30)), inference(subsumption_resolution, [], [f4355, f462])).
fof(f4355, plain, ((e1 = e4) | (~ spl4_27 | ~ spl4_30)), inference(forward_demodulation, [], [f735, f723])).
fof(f4353, plain, (~ spl4_32 | ~ spl4_34), inference(avatar_contradiction_clause, [], [f4352])).
fof(f4352, plain, ($false | (~ spl4_32 | ~ spl4_34)), inference(subsumption_resolution, [], [f4351, f461])).
fof(f4351, plain, ((e1 = e3) | (~ spl4_32 | ~ spl4_34)), inference(forward_demodulation, [], [f752, f744])).
fof(f4347, plain, (~ spl4_37 | ~ spl4_59 | spl4_83), inference(avatar_contradiction_clause, [], [f4346])).
fof(f4346, plain, ($false | (~ spl4_37 | ~ spl4_59 | spl4_83)), inference(subsumption_resolution, [], [f4343, f957])).
fof(f4343, plain, ((e2 = op(e1, e3)) | (~ spl4_37 | ~ spl4_59)), inference(backward_demodulation, [], [f4309, f765])).
fof(f4309, plain, ((e2 = op(op(e3, e2), e3)) | ~ spl4_59), inference(backward_demodulation, [], [f344, f857])).
fof(f4338, plain, (spl4_24 | ~ spl4_41 | ~ spl4_85), inference(avatar_split_clause, [], [f4334, f964, f780, f708])).
fof(f4334, plain, ((e3 = op(e4, e0)) | (~ spl4_41 | ~ spl4_85)), inference(backward_demodulation, [], [f4260, f782])).
fof(f4308, plain, (~ spl4_55 | ~ spl4_65), inference(avatar_split_clause, [], [f4298, f880, f838])).
fof(f4298, plain, (~ (e4 = op(e2, e4)) | ~ spl4_65), inference(backward_demodulation, [], [f434, f882])).
fof(f4306, plain, (~ spl4_15 | ~ spl4_65), inference(avatar_split_clause, [], [f4296, f880, f670])).
fof(f4296, plain, (~ (e4 = op(e4, e2)) | ~ spl4_65), inference(backward_demodulation, [], [f384, f882])).
fof(f4305, plain, (~ spl4_1 | ~ spl4_65), inference(avatar_contradiction_clause, [], [f4304])).
fof(f4304, plain, ($false | (~ spl4_1 | ~ spl4_65)), inference(subsumption_resolution, [], [f4303, f457])).
fof(f4303, plain, ((e0 = e2) | (~ spl4_1 | ~ spl4_65)), inference(forward_demodulation, [], [f4294, f614])).
fof(f4294, plain, ((e2 = op(e4, e4)) | ~ spl4_65), inference(backward_demodulation, [], [f343, f882])).
fof(f4283, plain, (~ spl4_17 | ~ spl4_77 | spl4_92), inference(avatar_contradiction_clause, [], [f4282])).
fof(f4282, plain, ($false | (~ spl4_17 | ~ spl4_77 | spl4_92)), inference(subsumption_resolution, [], [f4277, f995])).
fof(f995, plain, (~ (e1 = op(e1, e1)) | spl4_92), inference(avatar_component_clause, [], [f994])).
fof(f4277, plain, ((e1 = op(e1, e1)) | (~ spl4_17 | ~ spl4_77)), inference(backward_demodulation, [], [f4147, f933])).
fof(f4147, plain, ((e1 = op(e1, op(e1, e4))) | ~ spl4_17), inference(forward_demodulation, [], [f340, f681])).
fof(f4245, plain, (~ spl4_85 | ~ spl4_94 | spl4_146), inference(avatar_split_clause, [], [f4238, f1300, f1002, f964])).
fof(f4238, plain, (~ (e4 = op(e1, e3)) | (~ spl4_94 | spl4_146)), inference(backward_demodulation, [], [f1302, f1004])).
fof(f4231, plain, (~ spl4_98 | ~ spl4_99), inference(avatar_contradiction_clause, [], [f4230])).
fof(f4230, plain, ($false | (~ spl4_98 | ~ spl4_99)), inference(subsumption_resolution, [], [f4229, f463])).
fof(f4229, plain, ((e2 = e3) | (~ spl4_98 | ~ spl4_99)), inference(backward_demodulation, [], [f1025, f1021])).
fof(f4212, plain, (~ spl4_102 | ~ spl4_104), inference(avatar_contradiction_clause, [], [f4211])).
fof(f4211, plain, ($false | (~ spl4_102 | ~ spl4_104)), inference(subsumption_resolution, [], [f4210, f461])).
fof(f4210, plain, ((e1 = e3) | (~ spl4_102 | ~ spl4_104)), inference(backward_demodulation, [], [f1046, f1038])).
fof(f4194, plain, (spl4_104 | ~ spl4_50 | ~ spl4_106), inference(avatar_split_clause, [], [f4190, f1053, f817, f1044])).
fof(f4190, plain, ((e3 = op(e0, e4)) | (~ spl4_50 | ~ spl4_106)), inference(backward_demodulation, [], [f3172, f1055])).
fof(f4193, plain, (spl4_21 | ~ spl4_50 | ~ spl4_106), inference(avatar_split_clause, [], [f4189, f1053, f817, f696])).
fof(f4189, plain, ((e0 = op(e4, e0)) | (~ spl4_50 | ~ spl4_106)), inference(backward_demodulation, [], [f3171, f1055])).
fof(f3171, plain, ((e0 = op(e4, op(e0, e3))) | ~ spl4_50), inference(forward_demodulation, [], [f334, f819])).
fof(f4177, plain, (~ spl4_110 | ~ spl4_120), inference(avatar_split_clause, [], [f4175, f1111, f1069])).
fof(f4175, plain, (~ (e4 = op(e0, e3)) | ~ spl4_120), inference(backward_demodulation, [], [f411, f1113])).
fof(f4170, plain, (~ spl4_110 | ~ spl4_125), inference(avatar_split_clause, [], [f4163, f1132, f1069])).
fof(f4163, plain, (~ (e4 = op(e0, e3)) | ~ spl4_125), inference(backward_demodulation, [], [f408, f1134])).
fof(f4141, plain, (~ spl4_52 | ~ spl4_27), inference(avatar_split_clause, [], [f4140, f721, f826])).
fof(f4140, plain, (~ (e1 = op(e2, e4)) | ~ spl4_27), inference(forward_demodulation, [], [f403, f723])).
fof(f4133, plain, (spl4_94 | ~ spl4_32), inference(avatar_split_clause, [], [f4132, f742, f1002])).
fof(f4132, plain, ((e3 = op(e1, e1)) | ~ spl4_32), inference(forward_demodulation, [], [f349, f744])).
fof(f4121, plain, (~ spl4_37 | ~ spl4_32), inference(avatar_split_clause, [], [f4120, f742, f763])).
fof(f4120, plain, (~ (e1 = op(e3, e2)) | ~ spl4_32), inference(forward_demodulation, [], [f443, f744])).
fof(f4100, plain, (spl4_125 | ~ spl4_1), inference(avatar_split_clause, [], [f4099, f612, f1132])).
fof(f4099, plain, ((op(e0, e0) = e4) | ~ spl4_1), inference(forward_demodulation, [], [f355, f614])).
fof(f4075, plain, (~ spl4_4 | ~ spl4_32), inference(avatar_contradiction_clause, [], [f4074])).
fof(f4074, plain, ($false | (~ spl4_4 | ~ spl4_32)), inference(subsumption_resolution, [], [f4073, f462])).
fof(f4073, plain, ((e1 = e4) | (~ spl4_4 | ~ spl4_32)), inference(forward_demodulation, [], [f4065, f744])).
fof(f4065, plain, ((e4 = op(e3, e3)) | ~ spl4_4), inference(backward_demodulation, [], [f355, f626])).
fof(f4064, plain, (~ spl4_6 | ~ spl4_27 | ~ spl4_97), inference(avatar_contradiction_clause, [], [f4063])).
fof(f4063, plain, ($false | (~ spl4_6 | ~ spl4_27 | ~ spl4_97)), inference(subsumption_resolution, [], [f4062, f462])).
fof(f4062, plain, ((e1 = e4) | (~ spl4_6 | ~ spl4_27 | ~ spl4_97)), inference(forward_demodulation, [], [f4060, f1017])).
fof(f4060, plain, ((e4 = op(e1, e0)) | (~ spl4_6 | ~ spl4_27)), inference(backward_demodulation, [], [f4011, f635])).
fof(f4053, plain, (~ spl4_13 | ~ spl4_55), inference(avatar_contradiction_clause, [], [f4052])).
fof(f4052, plain, ($false | (~ spl4_13 | ~ spl4_55)), inference(subsumption_resolution, [], [f4051, f464])).
fof(f4051, plain, ((e2 = e4) | (~ spl4_13 | ~ spl4_55)), inference(forward_demodulation, [], [f4046, f664])).
fof(f4046, plain, ((e4 = op(e4, e2)) | (~ spl4_13 | ~ spl4_55)), inference(backward_demodulation, [], [f3969, f664])).
fof(f3969, plain, ((e4 = op(e4, op(e4, e2))) | ~ spl4_55), inference(backward_demodulation, [], [f353, f840])).
fof(f4040, plain, (~ spl4_17 | ~ spl4_20), inference(avatar_contradiction_clause, [], [f4039])).
fof(f4039, plain, ($false | (~ spl4_17 | ~ spl4_20)), inference(subsumption_resolution, [], [f4038, f462])).
fof(f4038, plain, ((e1 = e4) | (~ spl4_17 | ~ spl4_20)), inference(backward_demodulation, [], [f693, f681])).
fof(f4035, plain, (~ spl4_20 | ~ spl4_52 | ~ spl4_78), inference(avatar_contradiction_clause, [], [f4034])).
fof(f4034, plain, ($false | (~ spl4_20 | ~ spl4_52 | ~ spl4_78)), inference(subsumption_resolution, [], [f4033, f462])).
fof(f4033, plain, ((e1 = e4) | (~ spl4_20 | ~ spl4_52 | ~ spl4_78)), inference(forward_demodulation, [], [f4031, f828])).
fof(f4031, plain, ((e4 = op(e2, e4)) | (~ spl4_20 | ~ spl4_78)), inference(backward_demodulation, [], [f3947, f693])).
fof(f3947, plain, ((e4 = op(e2, op(e4, e1))) | ~ spl4_78), inference(backward_demodulation, [], [f352, f937])).
fof(f4009, plain, (~ spl4_32 | ~ spl4_33), inference(avatar_contradiction_clause, [], [f4008])).
fof(f4008, plain, ($false | (~ spl4_32 | ~ spl4_33)), inference(subsumption_resolution, [], [f4007, f460])).
fof(f4007, plain, ((e1 = e2) | (~ spl4_32 | ~ spl4_33)), inference(backward_demodulation, [], [f748, f744])).
fof(f3985, plain, (spl4_32 | ~ spl4_44 | ~ spl4_84), inference(avatar_split_clause, [], [f3981, f960, f792, f742])).
fof(f3981, plain, ((e1 = op(e3, e3)) | (~ spl4_44 | ~ spl4_84)), inference(backward_demodulation, [], [f3742, f794])).
fof(f3742, plain, ((e1 = op(op(e3, e1), e3)) | ~ spl4_84), inference(backward_demodulation, [], [f339, f962])).
fof(f3976, plain, (~ spl4_52 | ~ spl4_55), inference(avatar_contradiction_clause, [], [f3975])).
fof(f3975, plain, ($false | (~ spl4_52 | ~ spl4_55)), inference(subsumption_resolution, [], [f3974, f462])).
fof(f3974, plain, ((e1 = e4) | (~ spl4_52 | ~ spl4_55)), inference(backward_demodulation, [], [f840, f828])).
fof(f3944, plain, (spl4_42 | ~ spl4_97 | ~ spl4_119), inference(avatar_contradiction_clause, [], [f3943])).
fof(f3943, plain, ($false | (spl4_42 | ~ spl4_97 | ~ spl4_119)), inference(subsumption_resolution, [], [f3938, f785])).
fof(f785, plain, (~ (e1 = op(e3, e1)) | spl4_42), inference(avatar_component_clause, [], [f784])).
fof(f3938, plain, ((e1 = op(e3, e1)) | (~ spl4_97 | ~ spl4_119)), inference(backward_demodulation, [], [f3894, f1017])).
fof(f3894, plain, ((e1 = op(e3, op(e1, e0))) | ~ spl4_119), inference(backward_demodulation, [], [f336, f1109])).
fof(f3914, plain, (spl4_4 | ~ spl4_50 | ~ spl4_110), inference(avatar_split_clause, [], [f3908, f1069, f817, f624])).
fof(f3908, plain, ((e3 = op(e4, e4)) | (~ spl4_50 | ~ spl4_110)), inference(backward_demodulation, [], [f3172, f1071])).
fof(f3913, plain, (spl4_1 | ~ spl4_50 | ~ spl4_110), inference(avatar_split_clause, [], [f3907, f1069, f817, f612])).
fof(f3907, plain, ((e0 = op(e4, e4)) | (~ spl4_50 | ~ spl4_110)), inference(backward_demodulation, [], [f3171, f1071])).
fof(f3904, plain, (~ spl4_112 | ~ spl4_113), inference(avatar_contradiction_clause, [], [f3903])).
fof(f3903, plain, ($false | (~ spl4_112 | ~ spl4_113)), inference(subsumption_resolution, [], [f3902, f460])).
fof(f3902, plain, ((e1 = e2) | (~ spl4_112 | ~ spl4_113)), inference(backward_demodulation, [], [f1084, f1080])).
fof(f3849, plain, (spl4_122 | ~ spl4_91), inference(avatar_split_clause, [], [f3848, f990, f1120])).
fof(f3848, plain, ((op(e0, e0) = e1) | ~ spl4_91), inference(forward_demodulation, [], [f337, f992])).
fof(f3834, plain, (~ spl4_36 | ~ spl4_33 | spl4_169), inference(avatar_split_clause, [], [f3482, f1419, f746, f759])).
fof(f3482, plain, (~ (e0 = op(e3, e2)) | (~ spl4_33 | spl4_169)), inference(forward_demodulation, [], [f1421, f748])).
fof(f3832, plain, (~ spl4_14 | ~ spl4_64), inference(avatar_split_clause, [], [f3831, f876, f666])).
fof(f3831, plain, (~ (e3 = op(e4, e2)) | ~ spl4_64), inference(forward_demodulation, [], [f384, f878])).
fof(f3830, plain, (~ spl4_15 | ~ spl4_90), inference(avatar_split_clause, [], [f3829, f985, f670])).
fof(f3829, plain, (~ (e4 = op(e4, e2)) | ~ spl4_90), inference(forward_demodulation, [], [f382, f987])).
fof(f3820, plain, (~ spl4_7 | ~ spl4_19 | ~ spl4_80), inference(avatar_contradiction_clause, [], [f3819])).
fof(f3819, plain, ($false | (~ spl4_7 | ~ spl4_19 | ~ spl4_80)), inference(subsumption_resolution, [], [f3818, f462])).
fof(f3818, plain, ((e1 = e4) | (~ spl4_7 | ~ spl4_19 | ~ spl4_80)), inference(forward_demodulation, [], [f3815, f639])).
fof(f3815, plain, ((e4 = op(e4, e3)) | (~ spl4_19 | ~ spl4_80)), inference(backward_demodulation, [], [f3748, f689])).
fof(f3748, plain, ((e4 = op(e4, op(e4, e1))) | ~ spl4_80), inference(backward_demodulation, [], [f352, f945])).
fof(f3804, plain, (~ spl4_46 | ~ spl4_50), inference(avatar_contradiction_clause, [], [f3803])).
fof(f3803, plain, ($false | (~ spl4_46 | ~ spl4_50)), inference(subsumption_resolution, [], [f3802, f459])).
fof(f3802, plain, ((e0 = e4) | (~ spl4_46 | ~ spl4_50)), inference(backward_demodulation, [], [f819, f803])).
fof(f3786, plain, (~ spl4_37 | ~ spl4_60 | spl4_78), inference(avatar_contradiction_clause, [], [f3785])).
fof(f3785, plain, ($false | (~ spl4_37 | ~ spl4_60 | spl4_78)), inference(subsumption_resolution, [], [f3781, f936])).
fof(f936, plain, (~ (e2 = op(e1, e4)) | spl4_78), inference(avatar_component_clause, [], [f935])).
fof(f3781, plain, ((e2 = op(e1, e4)) | (~ spl4_37 | ~ spl4_60)), inference(backward_demodulation, [], [f3661, f861])).
fof(f3661, plain, ((e2 = op(e1, op(e2, e3))) | ~ spl4_37), inference(forward_demodulation, [], [f344, f765])).
fof(f3761, plain, (~ spl4_72 | spl4_86 | ~ spl4_113), inference(avatar_contradiction_clause, [], [f3760])).
fof(f3760, plain, ($false | (~ spl4_72 | spl4_86 | ~ spl4_113)), inference(subsumption_resolution, [], [f3757, f970])).
fof(f970, plain, (~ (e0 = op(e1, e2)) | spl4_86), inference(avatar_component_clause, [], [f969])).
fof(f3757, plain, ((e0 = op(e1, e2)) | (~ spl4_72 | ~ spl4_113)), inference(backward_demodulation, [], [f2640, f912])).
fof(f2640, plain, ((e0 = op(op(e2, e0), e2)) | ~ spl4_113), inference(forward_demodulation, [], [f333, f1084])).
fof(f3741, plain, (~ spl4_91 | ~ spl4_92), inference(avatar_contradiction_clause, [], [f3740])).
fof(f3740, plain, ($false | (~ spl4_91 | ~ spl4_92)), inference(subsumption_resolution, [], [f3739, f456])).
fof(f3739, plain, ((e0 = e1) | (~ spl4_91 | ~ spl4_92)), inference(backward_demodulation, [], [f996, f992])).
fof(f3737, plain, (~ spl4_11 | ~ spl4_98 | ~ spl4_120), inference(avatar_contradiction_clause, [], [f3736])).
fof(f3736, plain, ($false | (~ spl4_11 | ~ spl4_98 | ~ spl4_120)), inference(subsumption_resolution, [], [f3735, f456])).
fof(f3735, plain, ((e0 = e1) | (~ spl4_11 | ~ spl4_98 | ~ spl4_120)), inference(forward_demodulation, [], [f3731, f656])).
fof(f3731, plain, ((e1 = op(e4, e2)) | (~ spl4_98 | ~ spl4_120)), inference(backward_demodulation, [], [f3685, f1021])).
fof(f3685, plain, ((e1 = op(e4, op(e1, e0))) | ~ spl4_120), inference(backward_demodulation, [], [f336, f1113])).
fof(f3734, plain, (spl4_51 | ~ spl4_98 | ~ spl4_120), inference(avatar_split_clause, [], [f3730, f1111, f1019, f822])).
fof(f3730, plain, ((e0 = op(e2, e4)) | (~ spl4_98 | ~ spl4_120)), inference(backward_demodulation, [], [f3684, f1021])).
fof(f3684, plain, ((e0 = op(op(e1, e0), e4)) | ~ spl4_120), inference(backward_demodulation, [], [f332, f1113])).
fof(f3732, plain, (~ spl4_73 | ~ spl4_98), inference(avatar_split_clause, [], [f3724, f1019, f914])).
fof(f3724, plain, (~ (e2 = op(e2, e0)) | ~ spl4_98), inference(backward_demodulation, [], [f360, f1021])).
fof(f3711, plain, (~ spl4_56 | ~ spl4_106), inference(avatar_split_clause, [], [f3703, f1053, f843])).
fof(f3703, plain, (~ (e0 = op(e2, e3)) | ~ spl4_106), inference(backward_demodulation, [], [f387, f1055])).
fof(f3701, plain, (~ spl4_3 | ~ spl4_120 | ~ spl4_122 | spl4_178), inference(avatar_contradiction_clause, [], [f3700])).
fof(f3700, plain, ($false | (~ spl4_3 | ~ spl4_120 | ~ spl4_122 | spl4_178)), inference(subsumption_resolution, [], [f3693, f622])).
fof(f3693, plain, (~ (e2 = op(e4, e4)) | (~ spl4_120 | ~ spl4_122 | spl4_178)), inference(backward_demodulation, [], [f3464, f1113])).
fof(f3464, plain, (~ (e2 = op(op(e0, e1), op(e0, e1))) | (~ spl4_122 | spl4_178)), inference(forward_demodulation, [], [f1478, f1122])).
fof(f3683, plain, (~ spl4_122 | ~ spl4_124), inference(avatar_contradiction_clause, [], [f3682])).
fof(f3682, plain, ($false | (~ spl4_122 | ~ spl4_124)), inference(subsumption_resolution, [], [f3681, f461])).
fof(f3681, plain, ((e1 = e3) | (~ spl4_122 | ~ spl4_124)), inference(forward_demodulation, [], [f1130, f1122])).
fof(f3677, plain, (~ spl4_102 | ~ spl4_122), inference(avatar_split_clause, [], [f3466, f1120, f1036])).
fof(f3466, plain, (~ (e1 = op(e0, e4)) | ~ spl4_122), inference(forward_demodulation, [], [f409, f1122])).
fof(f3666, plain, (~ spl4_80 | ~ spl4_90), inference(avatar_split_clause, [], [f3665, f985, f943])).
fof(f3665, plain, (~ (e4 = op(e1, e4)) | ~ spl4_90), inference(forward_demodulation, [], [f424, f987])).
fof(f3660, plain, (~ spl4_59 | ~ spl4_64), inference(avatar_split_clause, [], [f3659, f876, f855])).
fof(f3659, plain, (~ (e3 = op(e2, e3)) | ~ spl4_64), inference(forward_demodulation, [], [f433, f878])).
fof(f3651, plain, (~ spl4_25 | ~ spl4_50), inference(avatar_split_clause, [], [f3650, f817, f712])).
fof(f3650, plain, (~ (e4 = op(e4, e0)) | ~ spl4_50), inference(forward_demodulation, [], [f365, f819])).
fof(f3649, plain, (~ spl4_23 | ~ spl4_3), inference(avatar_split_clause, [], [f3648, f620, f704])).
fof(f3648, plain, (~ (e2 = op(e4, e0)) | ~ spl4_3), inference(forward_demodulation, [], [f449, f622])).
fof(f3642, plain, (spl4_84 | ~ spl4_7 | ~ spl4_29), inference(avatar_split_clause, [], [f3641, f729, f637, f960])).
fof(f3641, plain, ((e3 = op(e1, e3)) | (~ spl4_7 | ~ spl4_29)), inference(forward_demodulation, [], [f3371, f639])).
fof(f3371, plain, ((e3 = op(op(e4, e3), e3)) | ~ spl4_29), inference(forward_demodulation, [], [f350, f731])).
fof(f3639, plain, (spl4_45 | ~ spl4_7 | ~ spl4_29), inference(avatar_split_clause, [], [f3638, f729, f637, f796])).
fof(f3638, plain, ((e4 = op(e3, e1)) | (~ spl4_7 | ~ spl4_29)), inference(forward_demodulation, [], [f3637, f731])).
fof(f3637, plain, ((e4 = op(op(e3, e4), e1)) | ~ spl4_7), inference(forward_demodulation, [], [f354, f639])).
fof(f3636, plain, (~ spl4_22 | ~ spl4_7), inference(avatar_split_clause, [], [f3635, f637, f700])).
fof(f3635, plain, (~ (e1 = op(e4, e0)) | ~ spl4_7), inference(forward_demodulation, [], [f448, f639])).
fof(f3629, plain, (~ spl4_3 | ~ spl4_64), inference(avatar_contradiction_clause, [], [f3628])).
fof(f3628, plain, ($false | (~ spl4_3 | ~ spl4_64)), inference(subsumption_resolution, [], [f3627, f465])).
fof(f3627, plain, ((e3 = e4) | (~ spl4_3 | ~ spl4_64)), inference(forward_demodulation, [], [f3626, f878])).
fof(f3626, plain, ((e4 = op(e2, e2)) | ~ spl4_3), inference(forward_demodulation, [], [f355, f622])).
fof(f3613, plain, (~ spl4_3 | ~ spl4_5), inference(avatar_contradiction_clause, [], [f3612])).
fof(f3612, plain, ($false | (~ spl4_3 | ~ spl4_5)), inference(subsumption_resolution, [], [f3611, f464])).
fof(f3611, plain, ((e2 = e4) | (~ spl4_3 | ~ spl4_5)), inference(backward_demodulation, [], [f630, f622])).
fof(f3609, plain, (spl4_3 | ~ spl4_5 | ~ spl4_175), inference(avatar_split_clause, [], [f3608, f1457, f628, f620])).
fof(f3608, plain, ((e2 = op(e4, e4)) | (~ spl4_5 | ~ spl4_175)), inference(forward_demodulation, [], [f3606, f630])).
fof(f3606, plain, ((e2 = op(op(e4, e4), op(e4, e4))) | (~ spl4_5 | ~ spl4_175)), inference(backward_demodulation, [], [f1458, f630])).
fof(f1458, plain, ((e2 = op(op(e4, op(e4, e4)), op(e4, op(e4, e4)))) | ~ spl4_175), inference(avatar_component_clause, [], [f1457])).
fof(f3601, plain, (~ spl4_7 | ~ spl4_9), inference(avatar_contradiction_clause, [], [f3600])).
fof(f3600, plain, ($false | (~ spl4_7 | ~ spl4_9)), inference(subsumption_resolution, [], [f3599, f461])).
fof(f3599, plain, ((e1 = e3) | (~ spl4_7 | ~ spl4_9)), inference(backward_demodulation, [], [f647, f639])).
fof(f3593, plain, (~ spl4_17 | ~ spl4_76 | spl4_97), inference(avatar_contradiction_clause, [], [f3592])).
fof(f3592, plain, ($false | (~ spl4_17 | ~ spl4_76 | spl4_97)), inference(subsumption_resolution, [], [f3589, f1016])).
fof(f3589, plain, ((e1 = op(e1, e0)) | (~ spl4_17 | ~ spl4_76)), inference(backward_demodulation, [], [f3387, f681])).
fof(f3387, plain, ((e1 = op(op(e4, e1), e0)) | ~ spl4_76), inference(forward_demodulation, [], [f340, f929])).
fof(f3591, plain, (~ spl4_2 | ~ spl4_17), inference(avatar_split_clause, [], [f3585, f679, f616])).
fof(f3585, plain, (~ (e1 = op(e4, e4)) | ~ spl4_17), inference(backward_demodulation, [], [f452, f681])).
fof(f3579, plain, (~ spl4_11 | ~ spl4_52 | ~ spl4_99), inference(avatar_contradiction_clause, [], [f3578])).
fof(f3578, plain, ($false | (~ spl4_11 | ~ spl4_52 | ~ spl4_99)), inference(subsumption_resolution, [], [f3577, f465])).
fof(f3577, plain, ((e3 = e4) | (~ spl4_11 | ~ spl4_52 | ~ spl4_99)), inference(forward_demodulation, [], [f3573, f1025])).
fof(f3573, plain, ((e4 = op(e1, e0)) | (~ spl4_11 | ~ spl4_52)), inference(backward_demodulation, [], [f3396, f828])).
fof(f3396, plain, ((e4 = op(op(e2, e4), e0)) | ~ spl4_11), inference(forward_demodulation, [], [f353, f656])).
fof(f3576, plain, (~ spl4_11 | ~ spl4_52 | spl4_118), inference(avatar_contradiction_clause, [], [f3575])).
fof(f3575, plain, ($false | (~ spl4_11 | ~ spl4_52 | spl4_118)), inference(subsumption_resolution, [], [f3572, f1104])).
fof(f1104, plain, (~ (e2 = op(e0, e1)) | spl4_118), inference(avatar_component_clause, [], [f1103])).
fof(f3572, plain, ((e2 = op(e0, e1)) | (~ spl4_11 | ~ spl4_52)), inference(backward_demodulation, [], [f3395, f828])).
fof(f3395, plain, ((e2 = op(e0, op(e2, e4))) | ~ spl4_11), inference(forward_demodulation, [], [f345, f656])).
fof(f3566, plain, (spl4_3 | ~ spl4_70 | ~ spl4_90), inference(avatar_split_clause, [], [f3563, f985, f901, f620])).
fof(f3563, plain, ((e2 = op(e4, e4)) | (~ spl4_70 | ~ spl4_90)), inference(backward_demodulation, [], [f3538, f903])).
fof(f3538, plain, ((e2 = op(e4, op(e2, e1))) | ~ spl4_90), inference(backward_demodulation, [], [f342, f987])).
fof(f3565, plain, (spl4_2 | ~ spl4_70 | ~ spl4_90), inference(avatar_split_clause, [], [f3562, f985, f901, f616])).
fof(f3562, plain, ((e1 = op(e4, e4)) | (~ spl4_70 | ~ spl4_90)), inference(backward_demodulation, [], [f3537, f903])).
fof(f3537, plain, ((e1 = op(op(e2, e1), e4)) | ~ spl4_90), inference(backward_demodulation, [], [f338, f987])).
fof(f3548, plain, (~ spl4_41 | ~ spl4_85 | ~ spl4_101), inference(avatar_contradiction_clause, [], [f3547])).
fof(f3547, plain, ($false | (~ spl4_41 | ~ spl4_85 | ~ spl4_101)), inference(subsumption_resolution, [], [f3546, f456])).
fof(f3546, plain, ((e0 = e1) | (~ spl4_41 | ~ spl4_85 | ~ spl4_101)), inference(forward_demodulation, [], [f3542, f1034])).
fof(f3542, plain, ((e1 = op(e0, e4)) | (~ spl4_41 | ~ spl4_85)), inference(backward_demodulation, [], [f3428, f966])).
fof(f3428, plain, ((e1 = op(e0, op(e1, e3))) | ~ spl4_41), inference(forward_demodulation, [], [f339, f782])).
fof(f3545, plain, (~ spl4_23 | ~ spl4_41 | ~ spl4_85), inference(avatar_contradiction_clause, [], [f3544])).
fof(f3544, plain, ($false | (~ spl4_23 | ~ spl4_41 | ~ spl4_85)), inference(subsumption_resolution, [], [f3543, f463])).
fof(f3543, plain, ((e2 = e3) | (~ spl4_23 | ~ spl4_41 | ~ spl4_85)), inference(forward_demodulation, [], [f3541, f706])).
fof(f3541, plain, ((e3 = op(e4, e0)) | (~ spl4_41 | ~ spl4_85)), inference(backward_demodulation, [], [f3429, f966])).
fof(f3429, plain, ((e3 = op(op(e1, e3), e0)) | ~ spl4_41), inference(forward_demodulation, [], [f347, f782])).
fof(f3519, plain, (spl4_4 | ~ spl4_50 | ~ spl4_110), inference(avatar_contradiction_clause, [], [f3518])).
fof(f3518, plain, ($false | (spl4_4 | ~ spl4_50 | ~ spl4_110)), inference(subsumption_resolution, [], [f3514, f625])).
fof(f625, plain, (~ (e3 = op(e4, e4)) | spl4_4), inference(avatar_component_clause, [], [f624])).
fof(f3514, plain, ((e3 = op(e4, e4)) | (~ spl4_50 | ~ spl4_110)), inference(backward_demodulation, [], [f3172, f1071])).
fof(f3515, plain, (~ spl4_85 | ~ spl4_110), inference(avatar_split_clause, [], [f3511, f1069, f964])).
fof(f3511, plain, (~ (e4 = op(e1, e3)) | ~ spl4_110), inference(backward_demodulation, [], [f386, f1071])).
fof(f3509, plain, (~ spl4_33 | ~ spl4_119 | ~ spl4_122 | spl4_178), inference(avatar_contradiction_clause, [], [f3508])).
fof(f3508, plain, ($false | (~ spl4_33 | ~ spl4_119 | ~ spl4_122 | spl4_178)), inference(subsumption_resolution, [], [f3497, f748])).
fof(f3497, plain, (~ (e2 = op(e3, e3)) | (~ spl4_119 | ~ spl4_122 | spl4_178)), inference(backward_demodulation, [], [f3464, f1109])).
fof(f3503, plain, (spl4_32 | ~ spl4_99 | ~ spl4_119), inference(avatar_contradiction_clause, [], [f3502])).
fof(f3502, plain, ($false | (spl4_32 | ~ spl4_99 | ~ spl4_119)), inference(subsumption_resolution, [], [f3493, f743])).
fof(f743, plain, (~ (e1 = op(e3, e3)) | spl4_32), inference(avatar_component_clause, [], [f742])).
fof(f3493, plain, ((e1 = op(e3, e3)) | (~ spl4_99 | ~ spl4_119)), inference(backward_demodulation, [], [f3238, f1109])).
fof(f3238, plain, ((e1 = op(op(e0, e1), e3)) | ~ spl4_99), inference(forward_demodulation, [], [f336, f1025])).
fof(f3441, plain, (~ spl4_17 | ~ spl4_92), inference(avatar_split_clause, [], [f3440, f994, f679])).
fof(f3440, plain, (~ (e1 = op(e4, e1)) | ~ spl4_92), inference(forward_demodulation, [], [f372, f996])).
fof(f3411, plain, (~ spl4_37 | ~ spl4_56 | spl4_98), inference(avatar_contradiction_clause, [], [f3410])).
fof(f3410, plain, ($false | (~ spl4_37 | ~ spl4_56 | spl4_98)), inference(subsumption_resolution, [], [f3409, f1020])).
fof(f1020, plain, (~ (e2 = op(e1, e0)) | spl4_98), inference(avatar_component_clause, [], [f1019])).
fof(f3409, plain, ((e2 = op(e1, e0)) | (~ spl4_37 | ~ spl4_56)), inference(forward_demodulation, [], [f3408, f765])).
fof(f3408, plain, ((e2 = op(op(e3, e2), e0)) | ~ spl4_56), inference(forward_demodulation, [], [f344, f845])).
fof(f3386, plain, (~ spl4_19 | ~ spl4_9), inference(avatar_split_clause, [], [f3385, f645, f687])).
fof(f3385, plain, (~ (e3 = op(e4, e1)) | ~ spl4_9), inference(forward_demodulation, [], [f451, f647])).
fof(f3382, plain, (~ spl4_9 | ~ spl4_29 | ~ spl4_33), inference(avatar_contradiction_clause, [], [f3381])).
fof(f3381, plain, ($false | (~ spl4_9 | ~ spl4_29 | ~ spl4_33)), inference(subsumption_resolution, [], [f3380, f464])).
fof(f3380, plain, ((e2 = e4) | (~ spl4_9 | ~ spl4_29 | ~ spl4_33)), inference(forward_demodulation, [], [f3379, f748])).
fof(f3379, plain, ((e4 = op(e3, e3)) | (~ spl4_9 | ~ spl4_29)), inference(forward_demodulation, [], [f3378, f731])).
fof(f3378, plain, ((e4 = op(op(e3, e4), e3)) | ~ spl4_9), inference(forward_demodulation, [], [f354, f647])).
fof(f3359, plain, (~ spl4_9 | ~ spl4_10), inference(avatar_contradiction_clause, [], [f3358])).
fof(f3358, plain, ($false | (~ spl4_9 | ~ spl4_10)), inference(subsumption_resolution, [], [f3357, f465])).
fof(f3357, plain, ((e3 = e4) | (~ spl4_9 | ~ spl4_10)), inference(backward_demodulation, [], [f651, f647])).
fof(f3355, plain, (~ spl4_10 | ~ spl4_29), inference(avatar_contradiction_clause, [], [f3354])).
fof(f3354, plain, ($false | (~ spl4_10 | ~ spl4_29)), inference(subsumption_resolution, [], [f3353, f465])).
fof(f3353, plain, ((e3 = e4) | (~ spl4_10 | ~ spl4_29)), inference(forward_demodulation, [], [f3351, f731])).
fof(f3351, plain, ((e4 = op(e3, e4)) | (~ spl4_10 | ~ spl4_29)), inference(backward_demodulation, [], [f2993, f731])).
fof(f2993, plain, ((e4 = op(op(e3, e4), e4)) | ~ spl4_10), inference(forward_demodulation, [], [f354, f651])).
fof(f3345, plain, (~ spl4_37 | ~ spl4_57 | ~ spl4_92), inference(avatar_contradiction_clause, [], [f3344])).
fof(f3344, plain, ($false | (~ spl4_37 | ~ spl4_57 | ~ spl4_92)), inference(subsumption_resolution, [], [f3343, f460])).
fof(f3343, plain, ((e1 = e2) | (~ spl4_37 | ~ spl4_57 | ~ spl4_92)), inference(forward_demodulation, [], [f3335, f996])).
fof(f3335, plain, ((e2 = op(e1, e1)) | (~ spl4_37 | ~ spl4_57)), inference(backward_demodulation, [], [f2852, f765])).
fof(f2852, plain, ((e2 = op(op(e3, e2), e1)) | ~ spl4_57), inference(forward_demodulation, [], [f344, f849])).
fof(f3311, plain, (spl4_32 | ~ spl4_69 | ~ spl4_89), inference(avatar_split_clause, [], [f3308, f981, f897, f742])).
fof(f3308, plain, ((e1 = op(e3, e3)) | (~ spl4_69 | ~ spl4_89)), inference(backward_demodulation, [], [f3273, f899])).
fof(f3273, plain, ((e1 = op(op(e2, e1), e3)) | ~ spl4_89), inference(backward_demodulation, [], [f338, f983])).
fof(f3293, plain, (~ spl4_26 | ~ spl4_76), inference(avatar_split_clause, [], [f3282, f927, f717])).
fof(f3282, plain, (~ (e0 = op(e3, e4)) | ~ spl4_76), inference(backward_demodulation, [], [f401, f929])).
fof(f3281, plain, (~ spl4_83 | ~ spl4_84), inference(avatar_contradiction_clause, [], [f3280])).
fof(f3280, plain, ($false | (~ spl4_83 | ~ spl4_84)), inference(subsumption_resolution, [], [f3279, f463])).
fof(f3279, plain, ((e2 = e3) | (~ spl4_83 | ~ spl4_84)), inference(backward_demodulation, [], [f962, f958])).
fof(f3272, plain, (~ spl4_92 | ~ spl4_95), inference(avatar_contradiction_clause, [], [f3271])).
fof(f3271, plain, ($false | (~ spl4_92 | ~ spl4_95)), inference(subsumption_resolution, [], [f3270, f462])).
fof(f3270, plain, ((e1 = e4) | (~ spl4_92 | ~ spl4_95)), inference(backward_demodulation, [], [f1008, f996])).
fof(f3267, plain, (spl4_16 | ~ spl4_50 | ~ spl4_107), inference(avatar_contradiction_clause, [], [f3266])).
fof(f3266, plain, ($false | (spl4_16 | ~ spl4_50 | ~ spl4_107)), inference(subsumption_resolution, [], [f3262, f676])).
fof(f676, plain, (~ (e0 = op(e4, e1)) | spl4_16), inference(avatar_component_clause, [], [f675])).
fof(f3262, plain, ((e0 = op(e4, e1)) | (~ spl4_50 | ~ spl4_107)), inference(backward_demodulation, [], [f3171, f1059])).
fof(f3265, plain, (~ spl4_50 | spl4_79 | ~ spl4_107), inference(avatar_contradiction_clause, [], [f3264])).
fof(f3264, plain, ($false | (~ spl4_50 | spl4_79 | ~ spl4_107)), inference(subsumption_resolution, [], [f3261, f940])).
fof(f3261, plain, ((e3 = op(e1, e4)) | (~ spl4_50 | ~ spl4_107)), inference(backward_demodulation, [], [f3172, f1059])).
fof(f3258, plain, (~ spl4_113 | ~ spl4_115), inference(avatar_contradiction_clause, [], [f3257])).
fof(f3257, plain, ($false | (~ spl4_113 | ~ spl4_115)), inference(subsumption_resolution, [], [f3256, f464])).
fof(f3256, plain, ((e2 = e4) | (~ spl4_113 | ~ spl4_115)), inference(forward_demodulation, [], [f1092, f1084])).
fof(f3248, plain, (~ spl4_121 | ~ spl4_122), inference(avatar_contradiction_clause, [], [f3247])).
fof(f3247, plain, ($false | (~ spl4_121 | ~ spl4_122)), inference(subsumption_resolution, [], [f3246, f456])).
fof(f3246, plain, ((e0 = e1) | (~ spl4_121 | ~ spl4_122)), inference(forward_demodulation, [], [f1122, f1118])).
fof(f3227, plain, (~ spl4_89 | ~ spl4_64), inference(avatar_split_clause, [], [f3226, f876, f981])).
fof(f3226, plain, (~ (e3 = op(e1, e2)) | ~ spl4_64), inference(forward_demodulation, [], [f380, f878])).
fof(f3205, plain, (~ spl4_74 | ~ spl4_64), inference(avatar_split_clause, [], [f3204, f876, f918])).
fof(f3204, plain, (~ (e3 = op(e2, e0)) | ~ spl4_64), inference(forward_demodulation, [], [f427, f878])).
fof(f3184, plain, (spl4_56 | ~ spl4_23 | ~ spl4_104), inference(avatar_split_clause, [], [f3145, f1044, f704, f843])).
fof(f3145, plain, ((e0 = op(e2, e3)) | (~ spl4_23 | ~ spl4_104)), inference(backward_demodulation, [], [f3044, f706])).
fof(f3044, plain, ((e0 = op(op(e4, e0), e3)) | ~ spl4_104), inference(backward_demodulation, [], [f335, f1046])).
fof(f3182, plain, (~ spl4_75 | ~ spl4_50), inference(avatar_split_clause, [], [f3181, f817, f922])).
fof(f3181, plain, (~ (e4 = op(e2, e0)) | ~ spl4_50), inference(forward_demodulation, [], [f363, f819])).
fof(f3168, plain, (spl4_40 | ~ spl4_23 | ~ spl4_104), inference(avatar_split_clause, [], [f3146, f1044, f704, f775])).
fof(f3146, plain, ((e4 = op(e3, e2)) | (~ spl4_23 | ~ spl4_104)), inference(backward_demodulation, [], [f3045, f706])).
fof(f3045, plain, ((e4 = op(e3, op(e4, e0))) | ~ spl4_104), inference(backward_demodulation, [], [f351, f1046])).
fof(f3114, plain, (~ spl4_47 | ~ spl4_50), inference(avatar_contradiction_clause, [], [f3113])).
fof(f3113, plain, ($false | (~ spl4_47 | ~ spl4_50)), inference(subsumption_resolution, [], [f3112, f462])).
fof(f3112, plain, ((e1 = e4) | (~ spl4_47 | ~ spl4_50)), inference(forward_demodulation, [], [f819, f807])).
fof(f3110, plain, (~ spl4_64 | ~ spl4_65), inference(avatar_contradiction_clause, [], [f3109])).
fof(f3109, plain, ($false | (~ spl4_64 | ~ spl4_65)), inference(subsumption_resolution, [], [f3108, f465])).
fof(f3108, plain, ((e3 = e4) | (~ spl4_64 | ~ spl4_65)), inference(backward_demodulation, [], [f882, f878])).
fof(f3098, plain, (~ spl4_2 | ~ spl4_65), inference(avatar_contradiction_clause, [], [f3097])).
fof(f3097, plain, ($false | (~ spl4_2 | ~ spl4_65)), inference(subsumption_resolution, [], [f3096, f460])).
fof(f3096, plain, ((e1 = e2) | (~ spl4_2 | ~ spl4_65)), inference(forward_demodulation, [], [f3090, f618])).
fof(f3090, plain, ((e2 = op(e4, e4)) | ~ spl4_65), inference(backward_demodulation, [], [f343, f882])).
fof(f3083, plain, (~ spl4_24 | ~ spl4_74), inference(avatar_split_clause, [], [f3078, f918, f708])).
fof(f3078, plain, (~ (e3 = op(e4, e0)) | ~ spl4_74), inference(backward_demodulation, [], [f364, f920])).
fof(f3068, plain, (~ spl4_81 | ~ spl4_84), inference(avatar_contradiction_clause, [], [f3067])).
fof(f3067, plain, ($false | (~ spl4_81 | ~ spl4_84)), inference(subsumption_resolution, [], [f3066, f458])).
fof(f3066, plain, ((e0 = e3) | (~ spl4_81 | ~ spl4_84)), inference(forward_demodulation, [], [f962, f950])).
fof(f3064, plain, (~ spl4_93 | ~ spl4_95), inference(avatar_contradiction_clause, [], [f3063])).
fof(f3063, plain, ($false | (~ spl4_93 | ~ spl4_95)), inference(subsumption_resolution, [], [f3062, f464])).
fof(f3062, plain, ((e2 = e4) | (~ spl4_93 | ~ spl4_95)), inference(backward_demodulation, [], [f1008, f1000])).
fof(f3043, plain, (~ spl4_105 | ~ spl4_110), inference(avatar_split_clause, [], [f3039, f1069, f1048])).
fof(f3039, plain, (~ (e4 = op(e0, e4)) | ~ spl4_110), inference(backward_demodulation, [], [f415, f1071])).
fof(f3042, plain, (spl4_19 | ~ spl4_47 | ~ spl4_110), inference(avatar_split_clause, [], [f3038, f1069, f805, f687])).
fof(f3038, plain, ((e3 = op(e4, e1)) | (~ spl4_47 | ~ spl4_110)), inference(backward_demodulation, [], [f3035, f1071])).
fof(f3035, plain, ((e3 = op(op(e0, e3), e1)) | ~ spl4_47), inference(forward_demodulation, [], [f346, f807])).
fof(f3041, plain, (~ spl4_47 | spl4_76 | ~ spl4_110), inference(avatar_contradiction_clause, [], [f3040])).
fof(f3040, plain, ($false | (~ spl4_47 | spl4_76 | ~ spl4_110)), inference(subsumption_resolution, [], [f3037, f928])).
fof(f928, plain, (~ (e0 = op(e1, e4)) | spl4_76), inference(avatar_component_clause, [], [f927])).
fof(f3037, plain, ((e0 = op(e1, e4)) | (~ spl4_47 | ~ spl4_110)), inference(backward_demodulation, [], [f3036, f1071])).
fof(f3036, plain, ((e0 = op(e1, op(e0, e3))) | ~ spl4_47), inference(forward_demodulation, [], [f334, f807])).
fof(f3029, plain, (~ spl4_108 | ~ spl4_113), inference(avatar_split_clause, [], [f3028, f1082, f1061])).
fof(f3028, plain, (~ (e2 = op(e0, e3)) | ~ spl4_113), inference(forward_demodulation, [], [f413, f1084])).
fof(f3023, plain, (~ spl4_100 | ~ spl4_95), inference(avatar_split_clause, [], [f3022, f1006, f1027])).
fof(f3022, plain, (~ (e4 = op(e1, e0)) | ~ spl4_95), inference(forward_demodulation, [], [f416, f1008])).
fof(f3018, plain, (~ spl4_80 | ~ spl4_95), inference(avatar_split_clause, [], [f2856, f1006, f943])).
fof(f2856, plain, (~ (e4 = op(e1, e4)) | ~ spl4_95), inference(forward_demodulation, [], [f422, f1008])).
fof(f3015, plain, (spl4_75 | ~ spl4_11 | ~ spl4_53), inference(avatar_split_clause, [], [f2802, f830, f654, f922])).
fof(f2802, plain, ((e4 = op(e2, e0)) | (~ spl4_11 | ~ spl4_53)), inference(forward_demodulation, [], [f2606, f656])).
fof(f2606, plain, ((e4 = op(e2, op(e4, e2))) | ~ spl4_53), inference(backward_demodulation, [], [f353, f832])).
fof(f3014, plain, (spl4_64 | ~ spl4_33), inference(avatar_split_clause, [], [f2935, f746, f876])).
fof(f2935, plain, ((e3 = op(e2, e2)) | ~ spl4_33), inference(backward_demodulation, [], [f349, f748])).
fof(f3000, plain, (~ spl4_32 | ~ spl4_47), inference(avatar_split_clause, [], [f2835, f805, f742])).
fof(f2835, plain, (~ (e1 = op(e3, e3)) | ~ spl4_47), inference(forward_demodulation, [], [f438, f807])).
fof(f2998, plain, (spl4_24 | ~ spl4_10 | ~ spl4_26), inference(avatar_split_clause, [], [f2997, f717, f649, f708])).
fof(f2997, plain, ((e3 = op(e4, e0)) | (~ spl4_10 | ~ spl4_26)), inference(forward_demodulation, [], [f2996, f651])).
fof(f2996, plain, ((e3 = op(op(e4, e3), e0)) | ~ spl4_26), inference(forward_demodulation, [], [f350, f719])).
fof(f2995, plain, (spl4_105 | ~ spl4_10 | ~ spl4_26), inference(avatar_split_clause, [], [f2994, f717, f649, f1048])).
fof(f2994, plain, ((e4 = op(e0, e4)) | (~ spl4_10 | ~ spl4_26)), inference(forward_demodulation, [], [f2993, f719])).
fof(f2952, plain, (~ spl4_2 | ~ spl4_33 | ~ spl4_40 | spl4_158), inference(avatar_contradiction_clause, [], [f2951])).
fof(f2951, plain, ($false | (~ spl4_2 | ~ spl4_33 | ~ spl4_40 | spl4_158)), inference(subsumption_resolution, [], [f2950, f618])).
fof(f2950, plain, (~ (e1 = op(e4, e4)) | (~ spl4_33 | ~ spl4_40 | spl4_158)), inference(forward_demodulation, [], [f2941, f777])).
fof(f2941, plain, (~ (e1 = op(op(e3, e2), op(e3, e2))) | (~ spl4_33 | spl4_158)), inference(backward_demodulation, [], [f1366, f748])).
fof(f2946, plain, (~ spl4_33 | ~ spl4_62), inference(avatar_contradiction_clause, [], [f2945])).
fof(f2945, plain, ($false | (~ spl4_33 | ~ spl4_62)), inference(subsumption_resolution, [], [f2944, f461])).
fof(f2944, plain, ((e1 = e3) | (~ spl4_33 | ~ spl4_62)), inference(forward_demodulation, [], [f2935, f870])).
fof(f2933, plain, (~ spl4_40 | ~ spl4_57 | ~ spl4_78), inference(avatar_contradiction_clause, [], [f2932])).
fof(f2932, plain, ($false | (~ spl4_40 | ~ spl4_57 | ~ spl4_78)), inference(subsumption_resolution, [], [f2931, f463])).
fof(f2931, plain, ((e2 = e3) | (~ spl4_40 | ~ spl4_57 | ~ spl4_78)), inference(forward_demodulation, [], [f2929, f937])).
fof(f2929, plain, ((e3 = op(e1, e4)) | (~ spl4_40 | ~ spl4_57)), inference(backward_demodulation, [], [f2811, f777])).
fof(f2811, plain, ((e3 = op(e1, op(e3, e2))) | ~ spl4_57), inference(forward_demodulation, [], [f348, f849])).
fof(f2914, plain, (~ spl4_62 | spl4_93), inference(avatar_contradiction_clause, [], [f2913])).
fof(f2913, plain, ($false | (~ spl4_62 | spl4_93)), inference(subsumption_resolution, [], [f2904, f999])).
fof(f999, plain, (~ (e2 = op(e1, e1)) | spl4_93), inference(avatar_component_clause, [], [f998])).
fof(f2904, plain, ((e2 = op(e1, e1)) | ~ spl4_62), inference(backward_demodulation, [], [f343, f870])).
fof(f2903, plain, (~ spl4_66 | ~ spl4_69), inference(avatar_contradiction_clause, [], [f2902])).
fof(f2902, plain, ($false | (~ spl4_66 | ~ spl4_69)), inference(subsumption_resolution, [], [f2901, f458])).
fof(f2901, plain, ((e0 = e3) | (~ spl4_66 | ~ spl4_69)), inference(forward_demodulation, [], [f899, f887])).
fof(f2899, plain, (~ spl4_64 | ~ spl4_78 | ~ spl4_95 | spl4_184), inference(avatar_split_clause, [], [f2896, f1518, f1006, f935, f876])).
fof(f2896, plain, (~ (e3 = op(e2, e2)) | (~ spl4_78 | ~ spl4_95 | spl4_184)), inference(backward_demodulation, [], [f2445, f937])).
fof(f2445, plain, (~ (e3 = op(op(e1, e4), op(e1, e4))) | (~ spl4_95 | spl4_184)), inference(backward_demodulation, [], [f1520, f1008])).
fof(f2898, plain, (~ spl4_28 | ~ spl4_78), inference(avatar_split_clause, [], [f2891, f935, f725])).
fof(f2891, plain, (~ (e2 = op(e3, e4)) | ~ spl4_78), inference(backward_demodulation, [], [f401, f937])).
fof(f2888, plain, (spl4_22 | ~ spl4_45 | ~ spl4_81), inference(avatar_contradiction_clause, [], [f2887])).
fof(f2887, plain, ($false | (spl4_22 | ~ spl4_45 | ~ spl4_81)), inference(subsumption_resolution, [], [f2881, f701])).
fof(f2881, plain, ((e1 = op(e4, e0)) | (~ spl4_45 | ~ spl4_81)), inference(backward_demodulation, [], [f2866, f950])).
fof(f2866, plain, ((e1 = op(e4, op(e1, e3))) | ~ spl4_45), inference(forward_demodulation, [], [f339, f798])).
fof(f2886, plain, (~ spl4_45 | ~ spl4_81 | spl4_104), inference(avatar_contradiction_clause, [], [f2885])).
fof(f2885, plain, ($false | (~ spl4_45 | ~ spl4_81 | spl4_104)), inference(subsumption_resolution, [], [f2880, f1045])).
fof(f1045, plain, (~ (e3 = op(e0, e4)) | spl4_104), inference(avatar_component_clause, [], [f1044])).
fof(f2880, plain, ((e3 = op(e0, e4)) | (~ spl4_45 | ~ spl4_81)), inference(backward_demodulation, [], [f2820, f950])).
fof(f2820, plain, ((e3 = op(op(e1, e3), e4)) | ~ spl4_45), inference(forward_demodulation, [], [f347, f798])).
fof(f2871, plain, (~ spl4_87 | ~ spl4_89), inference(avatar_contradiction_clause, [], [f2870])).
fof(f2870, plain, ($false | (~ spl4_87 | ~ spl4_89)), inference(subsumption_resolution, [], [f2869, f461])).
fof(f2869, plain, ((e1 = e3) | (~ spl4_87 | ~ spl4_89)), inference(forward_demodulation, [], [f983, f975])).
fof(f2842, plain, (~ spl4_57 | ~ spl4_60), inference(avatar_contradiction_clause, [], [f2841])).
fof(f2841, plain, ($false | (~ spl4_57 | ~ spl4_60)), inference(subsumption_resolution, [], [f2840, f462])).
fof(f2840, plain, ((e1 = e4) | (~ spl4_57 | ~ spl4_60)), inference(forward_demodulation, [], [f861, f849])).
fof(f2834, plain, (spl4_81 | ~ spl4_47 | ~ spl4_109), inference(avatar_split_clause, [], [f2833, f1065, f805, f948])).
fof(f2833, plain, ((e0 = op(e1, e3)) | (~ spl4_47 | ~ spl4_109)), inference(forward_demodulation, [], [f2177, f807])).
fof(f2177, plain, ((e0 = op(op(e3, e0), e3)) | ~ spl4_109), inference(backward_demodulation, [], [f334, f1067])).
fof(f2808, plain, (~ spl4_38 | ~ spl4_113), inference(avatar_split_clause, [], [f2807, f1082, f767])).
fof(f2807, plain, (~ (e2 = op(e3, e2)) | ~ spl4_113), inference(forward_demodulation, [], [f378, f1084])).
fof(f2800, plain, (~ spl4_6 | ~ spl4_11), inference(avatar_split_clause, [], [f2799, f654, f633])).
fof(f2799, plain, (~ (e0 = op(e4, e3)) | ~ spl4_11), inference(forward_demodulation, [], [f453, f656])).
fof(f2793, plain, (~ spl4_6 | ~ spl4_8), inference(avatar_contradiction_clause, [], [f2792])).
fof(f2792, plain, ($false | (~ spl4_6 | ~ spl4_8)), inference(subsumption_resolution, [], [f2791, f457])).
fof(f2791, plain, ((e0 = e2) | (~ spl4_6 | ~ spl4_8)), inference(forward_demodulation, [], [f643, f635])).
fof(f2790, plain, (~ spl4_11 | ~ spl4_15), inference(avatar_contradiction_clause, [], [f2789])).
fof(f2789, plain, ($false | (~ spl4_11 | ~ spl4_15)), inference(subsumption_resolution, [], [f2788, f459])).
fof(f2788, plain, ((e0 = e4) | (~ spl4_11 | ~ spl4_15)), inference(backward_demodulation, [], [f672, f656])).
fof(f2782, plain, (~ spl4_15 | ~ spl4_53), inference(avatar_contradiction_clause, [], [f2781])).
fof(f2781, plain, ($false | (~ spl4_15 | ~ spl4_53)), inference(subsumption_resolution, [], [f2780, f464])).
fof(f2780, plain, ((e2 = e4) | (~ spl4_15 | ~ spl4_53)), inference(forward_demodulation, [], [f2778, f832])).
fof(f2778, plain, ((e4 = op(e2, e4)) | (~ spl4_15 | ~ spl4_53)), inference(backward_demodulation, [], [f2606, f672])).
fof(f2769, plain, (~ spl4_2 | ~ spl4_19 | ~ spl4_31 | spl4_133), inference(avatar_contradiction_clause, [], [f2768])).
fof(f2768, plain, ($false | (~ spl4_2 | ~ spl4_19 | ~ spl4_31 | spl4_133)), inference(subsumption_resolution, [], [f2762, f740])).
fof(f2762, plain, (~ (e0 = op(e3, e3)) | (~ spl4_2 | ~ spl4_19 | spl4_133)), inference(backward_demodulation, [], [f1988, f689])).
fof(f1988, plain, (~ (e0 = op(op(e4, e1), op(e4, e1))) | (~ spl4_2 | spl4_133)), inference(backward_demodulation, [], [f1243, f618])).
fof(f2759, plain, (~ spl4_6 | ~ spl4_24 | ~ spl4_105), inference(avatar_contradiction_clause, [], [f2758])).
fof(f2758, plain, ($false | (~ spl4_6 | ~ spl4_24 | ~ spl4_105)), inference(subsumption_resolution, [], [f2757, f459])).
fof(f2757, plain, ((e0 = e4) | (~ spl4_6 | ~ spl4_24 | ~ spl4_105)), inference(forward_demodulation, [], [f2752, f635])).
fof(f2752, plain, ((e4 = op(e4, e3)) | (~ spl4_24 | ~ spl4_105)), inference(backward_demodulation, [], [f1659, f710])).
fof(f1659, plain, ((e4 = op(e4, op(e4, e0))) | ~ spl4_105), inference(backward_demodulation, [], [f351, f1050])).
fof(f2754, plain, (~ spl4_14 | ~ spl4_24), inference(avatar_split_clause, [], [f2749, f708, f666])).
fof(f2749, plain, (~ (e3 = op(e4, e2)) | ~ spl4_24), inference(backward_demodulation, [], [f447, f710])).
fof(f2734, plain, (~ spl4_31 | ~ spl4_121), inference(avatar_contradiction_clause, [], [f2733])).
fof(f2733, plain, ($false | (~ spl4_31 | ~ spl4_121)), inference(subsumption_resolution, [], [f2732, f458])).
fof(f2732, plain, ((e0 = e3) | (~ spl4_31 | ~ spl4_121)), inference(forward_demodulation, [], [f2722, f1118])).
fof(f2722, plain, ((op(e0, e0) = e3) | ~ spl4_31), inference(backward_demodulation, [], [f349, f740])).
fof(f2719, plain, (~ spl4_43 | ~ spl4_45), inference(avatar_contradiction_clause, [], [f2718])).
fof(f2718, plain, ($false | (~ spl4_43 | ~ spl4_45)), inference(subsumption_resolution, [], [f2717, f464])).
fof(f2717, plain, ((e2 = e4) | (~ spl4_43 | ~ spl4_45)), inference(forward_demodulation, [], [f798, f790])).
fof(f2716, plain, (~ spl4_47 | ~ spl4_49), inference(avatar_contradiction_clause, [], [f2715])).
fof(f2715, plain, ($false | (~ spl4_47 | ~ spl4_49)), inference(subsumption_resolution, [], [f2714, f461])).
fof(f2714, plain, ((e1 = e3) | (~ spl4_47 | ~ spl4_49)), inference(backward_demodulation, [], [f815, f807])).
fof(f2711, plain, (~ spl4_57 | ~ spl4_59), inference(avatar_contradiction_clause, [], [f2710])).
fof(f2710, plain, ($false | (~ spl4_57 | ~ spl4_59)), inference(subsumption_resolution, [], [f2709, f461])).
fof(f2709, plain, ((e1 = e3) | (~ spl4_57 | ~ spl4_59)), inference(backward_demodulation, [], [f857, f849])).
fof(f2676, plain, (~ spl4_76 | ~ spl4_79), inference(avatar_contradiction_clause, [], [f2675])).
fof(f2675, plain, ($false | (~ spl4_76 | ~ spl4_79)), inference(subsumption_resolution, [], [f2674, f458])).
fof(f2674, plain, ((e0 = e3) | (~ spl4_76 | ~ spl4_79)), inference(backward_demodulation, [], [f941, f929])).
fof(f2658, plain, (~ spl4_90 | ~ spl4_95), inference(avatar_split_clause, [], [f2657, f1006, f985])).
fof(f2657, plain, (~ (e4 = op(e1, e2)) | ~ spl4_95), inference(forward_demodulation, [], [f420, f1008])).
fof(f2656, plain, (~ spl4_88 | ~ spl4_98), inference(avatar_split_clause, [], [f2655, f1019, f977])).
fof(f2655, plain, (~ (e2 = op(e1, e2)) | ~ spl4_98), inference(forward_demodulation, [], [f417, f1021])).
fof(f2636, plain, (~ spl4_63 | ~ spl4_53), inference(avatar_split_clause, [], [f2635, f830, f872])).
fof(f2635, plain, (~ (e2 = op(e2, e2)) | ~ spl4_53), inference(forward_demodulation, [], [f434, f832])).
fof(f2634, plain, (~ spl4_59 | ~ spl4_60), inference(avatar_contradiction_clause, [], [f2633])).
fof(f2633, plain, ($false | (~ spl4_59 | ~ spl4_60)), inference(subsumption_resolution, [], [f2632, f465])).
fof(f2632, plain, ((e3 = e4) | (~ spl4_59 | ~ spl4_60)), inference(forward_demodulation, [], [f861, f857])).
fof(f2616, plain, (~ spl4_37 | ~ spl4_59 | ~ spl4_82), inference(avatar_contradiction_clause, [], [f2615])).
fof(f2615, plain, ($false | (~ spl4_37 | ~ spl4_59 | ~ spl4_82)), inference(subsumption_resolution, [], [f2614, f460])).
fof(f2614, plain, ((e1 = e2) | (~ spl4_37 | ~ spl4_59 | ~ spl4_82)), inference(forward_demodulation, [], [f2613, f954])).
fof(f2613, plain, ((e2 = op(e1, e3)) | (~ spl4_37 | ~ spl4_59)), inference(forward_demodulation, [], [f2612, f765])).
fof(f2612, plain, ((e2 = op(op(e3, e2), e3)) | ~ spl4_59), inference(forward_demodulation, [], [f344, f857])).
fof(f2605, plain, (spl4_34 | ~ spl4_49 | ~ spl4_109), inference(avatar_contradiction_clause, [], [f2604])).
fof(f2604, plain, ($false | (spl4_34 | ~ spl4_49 | ~ spl4_109)), inference(subsumption_resolution, [], [f2603, f751])).
fof(f751, plain, (~ (e3 = op(e3, e3)) | spl4_34), inference(avatar_component_clause, [], [f750])).
fof(f2603, plain, ((e3 = op(e3, e3)) | (~ spl4_49 | ~ spl4_109)), inference(forward_demodulation, [], [f2178, f815])).
fof(f2178, plain, ((e3 = op(e3, op(e3, e0))) | ~ spl4_109), inference(backward_demodulation, [], [f346, f1067])).
fof(f2585, plain, (spl4_89 | ~ spl4_43 | ~ spl4_82), inference(avatar_split_clause, [], [f2584, f952, f788, f981])).
fof(f2584, plain, ((e3 = op(e1, e2)) | (~ spl4_43 | ~ spl4_82)), inference(forward_demodulation, [], [f2583, f954])).
fof(f2583, plain, ((e3 = op(op(e1, e3), e2)) | ~ spl4_43), inference(forward_demodulation, [], [f347, f790])).
fof(f2582, plain, (~ spl4_68 | ~ spl4_43), inference(avatar_split_clause, [], [f2581, f788, f893])).
fof(f2581, plain, (~ (e2 = op(e2, e1)) | ~ spl4_43), inference(forward_demodulation, [], [f373, f790])).
fof(f2564, plain, (~ spl4_31 | ~ spl4_79 | ~ spl4_95 | spl4_144), inference(avatar_split_clause, [], [f2448, f1292, f1006, f939, f738])).
fof(f2448, plain, (~ (e0 = op(e3, e3)) | (~ spl4_79 | ~ spl4_95 | spl4_144)), inference(forward_demodulation, [], [f2443, f941])).
fof(f2443, plain, (~ (e0 = op(op(e1, e4), op(e1, e4))) | (~ spl4_95 | spl4_144)), inference(backward_demodulation, [], [f1294, f1008])).
fof(f2519, plain, (~ spl4_2 | ~ spl4_35), inference(avatar_contradiction_clause, [], [f2518])).
fof(f2518, plain, ($false | (~ spl4_2 | ~ spl4_35)), inference(subsumption_resolution, [], [f2517, f461])).
fof(f2517, plain, ((e1 = e3) | (~ spl4_2 | ~ spl4_35)), inference(forward_demodulation, [], [f2510, f618])).
fof(f2510, plain, ((e3 = op(e4, e4)) | ~ spl4_35), inference(backward_demodulation, [], [f349, f756])).
fof(f2509, plain, (~ spl4_37 | ~ spl4_40), inference(avatar_contradiction_clause, [], [f2508])).
fof(f2508, plain, ($false | (~ spl4_37 | ~ spl4_40)), inference(subsumption_resolution, [], [f2507, f462])).
fof(f2507, plain, ((e1 = e4) | (~ spl4_37 | ~ spl4_40)), inference(backward_demodulation, [], [f777, f765])).
fof(f2498, plain, (~ spl4_49 | ~ spl4_50), inference(avatar_contradiction_clause, [], [f2497])).
fof(f2497, plain, ($false | (~ spl4_49 | ~ spl4_50)), inference(subsumption_resolution, [], [f2496, f465])).
fof(f2496, plain, ((e3 = e4) | (~ spl4_49 | ~ spl4_50)), inference(backward_demodulation, [], [f819, f815])).
fof(f2482, plain, (~ spl4_58 | ~ spl4_59), inference(avatar_contradiction_clause, [], [f2481])).
fof(f2481, plain, ($false | (~ spl4_58 | ~ spl4_59)), inference(subsumption_resolution, [], [f2480, f463])).
fof(f2480, plain, ((e2 = e3) | (~ spl4_58 | ~ spl4_59)), inference(forward_demodulation, [], [f857, f853])).
fof(f2458, plain, (~ spl4_82 | ~ spl4_85), inference(avatar_contradiction_clause, [], [f2457])).
fof(f2457, plain, ($false | (~ spl4_82 | ~ spl4_85)), inference(subsumption_resolution, [], [f2456, f462])).
fof(f2456, plain, ((e1 = e4) | (~ spl4_82 | ~ spl4_85)), inference(backward_demodulation, [], [f966, f954])).
fof(f2436, plain, (~ spl4_86 | ~ spl4_98 | ~ spl4_117), inference(avatar_contradiction_clause, [], [f2435])).
fof(f2435, plain, ($false | (~ spl4_86 | ~ spl4_98 | ~ spl4_117)), inference(subsumption_resolution, [], [f2434, f456])).
fof(f2434, plain, ((e0 = e1) | (~ spl4_86 | ~ spl4_98 | ~ spl4_117)), inference(forward_demodulation, [], [f2430, f971])).
fof(f2430, plain, ((e1 = op(e1, e2)) | (~ spl4_98 | ~ spl4_117)), inference(backward_demodulation, [], [f2162, f1021])).
fof(f2162, plain, ((e1 = op(e1, op(e1, e0))) | ~ spl4_117), inference(backward_demodulation, [], [f336, f1101])).
fof(f2398, plain, (spl4_107 | ~ spl4_16 | ~ spl4_79), inference(avatar_split_clause, [], [f2392, f939, f675, f1057])).
fof(f2392, plain, ((e1 = op(e0, e3)) | (~ spl4_16 | ~ spl4_79)), inference(backward_demodulation, [], [f2383, f941])).
fof(f2383, plain, ((e1 = op(e0, op(e1, e4))) | ~ spl4_16), inference(forward_demodulation, [], [f340, f677])).
fof(f2397, plain, (~ spl4_84 | ~ spl4_79), inference(avatar_split_clause, [], [f2389, f939, f960])).
fof(f2389, plain, (~ (e3 = op(e1, e3)) | ~ spl4_79), inference(backward_demodulation, [], [f425, f941])).
fof(f2378, plain, (spl4_54 | ~ spl4_40 | ~ spl4_58), inference(avatar_split_clause, [], [f2377, f851, f775, f834])).
fof(f2377, plain, ((e3 = op(e2, e4)) | (~ spl4_40 | ~ spl4_58)), inference(forward_demodulation, [], [f2376, f853])).
fof(f2376, plain, ((e3 = op(op(e2, e3), e4)) | ~ spl4_40), inference(forward_demodulation, [], [f348, f777])).
fof(f2369, plain, (~ spl4_68 | ~ spl4_58), inference(avatar_split_clause, [], [f2368, f851, f893])).
fof(f2368, plain, (~ (e2 = op(e2, e1)) | ~ spl4_58), inference(forward_demodulation, [], [f431, f853])).
fof(f2361, plain, (~ spl4_58 | ~ spl4_60), inference(avatar_contradiction_clause, [], [f2360])).
fof(f2360, plain, ($false | (~ spl4_58 | ~ spl4_60)), inference(subsumption_resolution, [], [f2359, f464])).
fof(f2359, plain, ((e2 = e4) | (~ spl4_58 | ~ spl4_60)), inference(forward_demodulation, [], [f861, f853])).
fof(f2323, plain, (~ spl4_115 | ~ spl4_40), inference(avatar_split_clause, [], [f2322, f775, f1090])).
fof(f2322, plain, (~ (e4 = op(e0, e2)) | ~ spl4_40), inference(forward_demodulation, [], [f378, f777])).
fof(f2321, plain, (~ spl4_34 | ~ spl4_109), inference(avatar_split_clause, [], [f2320, f1065, f750])).
fof(f2320, plain, (~ (e3 = op(e3, e3)) | ~ spl4_109), inference(forward_demodulation, [], [f388, f1067])).
fof(f2307, plain, (~ spl4_16 | ~ spl4_20), inference(avatar_contradiction_clause, [], [f2306])).
fof(f2306, plain, ($false | (~ spl4_16 | ~ spl4_20)), inference(subsumption_resolution, [], [f2305, f459])).
fof(f2305, plain, ((e0 = e4) | (~ spl4_16 | ~ spl4_20)), inference(forward_demodulation, [], [f693, f677])).
fof(f2290, plain, (~ spl4_16 | ~ spl4_18), inference(avatar_contradiction_clause, [], [f2289])).
fof(f2289, plain, ($false | (~ spl4_16 | ~ spl4_18)), inference(subsumption_resolution, [], [f2288, f457])).
fof(f2288, plain, ((e0 = e2) | (~ spl4_16 | ~ spl4_18)), inference(backward_demodulation, [], [f685, f677])).
fof(f2276, plain, (~ spl4_42 | ~ spl4_44), inference(avatar_contradiction_clause, [], [f2275])).
fof(f2275, plain, ($false | (~ spl4_42 | ~ spl4_44)), inference(subsumption_resolution, [], [f2274, f461])).
fof(f2274, plain, ((e1 = e3) | (~ spl4_42 | ~ spl4_44)), inference(forward_demodulation, [], [f794, f786])).
fof(f2273, plain, (~ spl4_48 | ~ spl4_50), inference(avatar_contradiction_clause, [], [f2272])).
fof(f2272, plain, ($false | (~ spl4_48 | ~ spl4_50)), inference(subsumption_resolution, [], [f2271, f464])).
fof(f2271, plain, ((e2 = e4) | (~ spl4_48 | ~ spl4_50)), inference(backward_demodulation, [], [f819, f811])).
fof(f2259, plain, (~ spl4_56 | ~ spl4_60), inference(avatar_contradiction_clause, [], [f2258])).
fof(f2258, plain, ($false | (~ spl4_56 | ~ spl4_60)), inference(subsumption_resolution, [], [f2257, f459])).
fof(f2257, plain, ((e0 = e4) | (~ spl4_56 | ~ spl4_60)), inference(forward_demodulation, [], [f861, f845])).
fof(f2248, plain, (~ spl4_62 | ~ spl4_92), inference(avatar_contradiction_clause, [], [f2247])).
fof(f2247, plain, ($false | (~ spl4_62 | ~ spl4_92)), inference(subsumption_resolution, [], [f2246, f460])).
fof(f2246, plain, ((e1 = e2) | (~ spl4_62 | ~ spl4_92)), inference(forward_demodulation, [], [f2241, f996])).
fof(f2241, plain, ((e2 = op(e1, e1)) | ~ spl4_62), inference(backward_demodulation, [], [f343, f870])).
fof(f2230, plain, (~ spl4_72 | spl4_91 | ~ spl4_112), inference(avatar_contradiction_clause, [], [f2229])).
fof(f2229, plain, ($false | (~ spl4_72 | spl4_91 | ~ spl4_112)), inference(subsumption_resolution, [], [f2222, f991])).
fof(f2222, plain, ((e0 = op(e1, e1)) | (~ spl4_72 | ~ spl4_112)), inference(backward_demodulation, [], [f2171, f912])).
fof(f2171, plain, ((e0 = op(op(e2, e0), e1)) | ~ spl4_112), inference(backward_demodulation, [], [f333, f1080])).
fof(f2218, plain, (~ spl4_83 | ~ spl4_85), inference(avatar_contradiction_clause, [], [f2217])).
fof(f2217, plain, ($false | (~ spl4_83 | ~ spl4_85)), inference(subsumption_resolution, [], [f2216, f464])).
fof(f2216, plain, ((e2 = e4) | (~ spl4_83 | ~ spl4_85)), inference(backward_demodulation, [], [f966, f958])).
fof(f2189, plain, (spl4_91 | ~ spl4_97 | ~ spl4_117), inference(avatar_split_clause, [], [f2183, f1099, f1015, f990])).
fof(f2183, plain, ((e0 = op(e1, e1)) | (~ spl4_97 | ~ spl4_117)), inference(backward_demodulation, [], [f2161, f1017])).
fof(f2161, plain, ((e0 = op(op(e1, e0), e1)) | ~ spl4_117), inference(backward_demodulation, [], [f332, f1101])).
fof(f2188, plain, (spl4_92 | ~ spl4_97 | ~ spl4_117), inference(avatar_split_clause, [], [f2182, f1099, f1015, f994])).
fof(f2182, plain, ((e1 = op(e1, e1)) | (~ spl4_97 | ~ spl4_117)), inference(backward_demodulation, [], [f2162, f1017])).
fof(f2170, plain, (~ spl4_107 | ~ spl4_117), inference(avatar_split_clause, [], [f2166, f1099, f1057])).
fof(f2166, plain, (~ (e1 = op(e0, e3)) | ~ spl4_117), inference(backward_demodulation, [], [f411, f1101])).
fof(f2168, plain, (~ spl4_67 | ~ spl4_117), inference(avatar_split_clause, [], [f2164, f1099, f889])).
fof(f2164, plain, (~ (e1 = op(e2, e1)) | ~ spl4_117), inference(backward_demodulation, [], [f367, f1101])).
fof(f2160, plain, (~ spl4_117 | ~ spl4_42), inference(avatar_split_clause, [], [f2159, f784, f1099])).
fof(f2159, plain, (~ (e1 = op(e0, e1)) | ~ spl4_42), inference(forward_demodulation, [], [f368, f786])).
fof(f2146, plain, (spl4_114 | ~ spl4_38 | ~ spl4_56), inference(avatar_split_clause, [], [f1882, f843, f767, f1086])).
fof(f1882, plain, ((e3 = op(e0, e2)) | (~ spl4_38 | ~ spl4_56)), inference(backward_demodulation, [], [f1825, f769])).
fof(f1825, plain, ((e3 = op(e0, op(e3, e2))) | ~ spl4_56), inference(backward_demodulation, [], [f348, f845])).
fof(f2137, plain, (~ spl4_103 | ~ spl4_105), inference(avatar_contradiction_clause, [], [f2136])).
fof(f2136, plain, ($false | (~ spl4_103 | ~ spl4_105)), inference(subsumption_resolution, [], [f2135, f464])).
fof(f2135, plain, ((e2 = e4) | (~ spl4_103 | ~ spl4_105)), inference(forward_demodulation, [], [f1050, f1042])).
fof(f2134, plain, (~ spl4_97 | ~ spl4_22), inference(avatar_split_clause, [], [f2133, f700, f1015])).
fof(f2133, plain, (~ (e1 = op(e1, e0)) | ~ spl4_22), inference(forward_demodulation, [], [f362, f702])).
fof(f2130, plain, (~ spl4_97 | ~ spl4_77), inference(avatar_split_clause, [], [f2129, f931, f1015])).
fof(f2129, plain, (~ (e1 = op(e1, e0)) | ~ spl4_77), inference(forward_demodulation, [], [f419, f933])).
fof(f2120, plain, (spl4_95 | ~ spl4_2), inference(avatar_split_clause, [], [f1984, f616, f1006])).
fof(f1984, plain, ((e4 = op(e1, e1)) | ~ spl4_2), inference(backward_demodulation, [], [f355, f618])).
fof(f2109, plain, (spl4_67 | ~ spl4_18 | ~ spl4_77), inference(avatar_split_clause, [], [f2108, f931, f683, f889])).
fof(f2108, plain, ((e1 = op(e2, e1)) | (~ spl4_18 | ~ spl4_77)), inference(forward_demodulation, [], [f2107, f685])).
fof(f2107, plain, ((e1 = op(op(e4, e1), e1)) | ~ spl4_77), inference(forward_demodulation, [], [f340, f933])).
fof(f2093, plain, (~ spl4_77 | ~ spl4_79), inference(avatar_contradiction_clause, [], [f2092])).
fof(f2092, plain, ($false | (~ spl4_77 | ~ spl4_79)), inference(subsumption_resolution, [], [f2091, f461])).
fof(f2091, plain, ((e1 = e3) | (~ spl4_77 | ~ spl4_79)), inference(forward_demodulation, [], [f941, f933])).
fof(f2069, plain, (~ spl4_62 | ~ spl4_52), inference(avatar_split_clause, [], [f2068, f826, f868])).
fof(f2068, plain, (~ (e1 = op(e2, e2)) | ~ spl4_52), inference(forward_demodulation, [], [f434, f828])).
fof(f2067, plain, (~ spl4_61 | ~ spl4_2 | ~ spl4_18 | spl4_133), inference(avatar_split_clause, [], [f1994, f1241, f683, f616, f864])).
fof(f1994, plain, (~ (e0 = op(e2, e2)) | (~ spl4_2 | ~ spl4_18 | spl4_133)), inference(forward_demodulation, [], [f1988, f685])).
fof(f2066, plain, (~ spl4_63 | ~ spl4_2 | ~ spl4_18 | spl4_175), inference(avatar_split_clause, [], [f1997, f1457, f683, f616, f872])).
fof(f1997, plain, (~ (e2 = op(e2, e2)) | (~ spl4_2 | ~ spl4_18 | spl4_175)), inference(forward_demodulation, [], [f1989, f685])).
fof(f1989, plain, (~ (e2 = op(op(e4, e1), op(e4, e1))) | (~ spl4_2 | spl4_175)), inference(backward_demodulation, [], [f1459, f618])).
fof(f2059, plain, (spl4_85 | ~ spl4_14 | ~ spl4_52), inference(avatar_split_clause, [], [f2058, f826, f666, f964])).
fof(f2058, plain, ((e4 = op(e1, e3)) | (~ spl4_14 | ~ spl4_52)), inference(forward_demodulation, [], [f2057, f828])).
fof(f2057, plain, ((e4 = op(op(e2, e4), e3)) | ~ spl4_14), inference(forward_demodulation, [], [f353, f668])).
fof(f2051, plain, (~ spl4_52 | ~ spl4_54), inference(avatar_contradiction_clause, [], [f2050])).
fof(f2050, plain, ($false | (~ spl4_52 | ~ spl4_54)), inference(subsumption_resolution, [], [f2049, f461])).
fof(f2049, plain, ((e1 = e3) | (~ spl4_52 | ~ spl4_54)), inference(forward_demodulation, [], [f836, f828])).
fof(f2038, plain, (spl4_99 | ~ spl4_7 | ~ spl4_26), inference(avatar_split_clause, [], [f2037, f717, f637, f1023])).
fof(f2037, plain, ((e3 = op(e1, e0)) | (~ spl4_7 | ~ spl4_26)), inference(forward_demodulation, [], [f2036, f639])).
fof(f2036, plain, ((e3 = op(op(e4, e3), e0)) | ~ spl4_26), inference(forward_demodulation, [], [f350, f719])).
fof(f2035, plain, (spl4_120 | ~ spl4_7 | ~ spl4_26), inference(avatar_split_clause, [], [f2034, f717, f637, f1111])).
fof(f2034, plain, ((e4 = op(e0, e1)) | (~ spl4_7 | ~ spl4_26)), inference(forward_demodulation, [], [f2033, f719])).
fof(f2033, plain, ((e4 = op(op(e3, e4), e1)) | ~ spl4_7), inference(forward_demodulation, [], [f354, f639])).
fof(f2020, plain, (spl4_76 | ~ spl4_22 | ~ spl4_105), inference(avatar_split_clause, [], [f2019, f1048, f700, f927])).
fof(f2019, plain, ((e0 = op(e1, e4)) | (~ spl4_22 | ~ spl4_105)), inference(forward_demodulation, [], [f1658, f702])).
fof(f1658, plain, ((e0 = op(op(e4, e0), e4)) | ~ spl4_105), inference(backward_demodulation, [], [f335, f1050])).
fof(f2014, plain, (~ spl4_22 | ~ spl4_24), inference(avatar_contradiction_clause, [], [f2013])).
fof(f2013, plain, ($false | (~ spl4_22 | ~ spl4_24)), inference(subsumption_resolution, [], [f2012, f461])).
fof(f2012, plain, ((e1 = e3) | (~ spl4_22 | ~ spl4_24)), inference(forward_demodulation, [], [f710, f702])).
fof(f2011, plain, (~ spl4_18 | ~ spl4_20), inference(avatar_contradiction_clause, [], [f2010])).
fof(f2010, plain, ($false | (~ spl4_18 | ~ spl4_20)), inference(subsumption_resolution, [], [f2009, f464])).
fof(f2009, plain, ((e2 = e4) | (~ spl4_18 | ~ spl4_20)), inference(forward_demodulation, [], [f693, f685])).
fof(f1993, plain, (~ spl4_2 | ~ spl4_91), inference(avatar_contradiction_clause, [], [f1992])).
fof(f1992, plain, ($false | (~ spl4_2 | ~ spl4_91)), inference(subsumption_resolution, [], [f1991, f459])).
fof(f1991, plain, ((e0 = e4) | (~ spl4_2 | ~ spl4_91)), inference(forward_demodulation, [], [f1984, f992])).
fof(f1983, plain, (~ spl4_6 | ~ spl4_7), inference(avatar_contradiction_clause, [], [f1982])).
fof(f1982, plain, ($false | (~ spl4_6 | ~ spl4_7)), inference(subsumption_resolution, [], [f1981, f456])).
fof(f1981, plain, ((e0 = e1) | (~ spl4_6 | ~ spl4_7)), inference(backward_demodulation, [], [f639, f635])).
fof(f1980, plain, (~ spl4_7 | ~ spl4_8), inference(avatar_contradiction_clause, [], [f1979])).
fof(f1979, plain, ($false | (~ spl4_7 | ~ spl4_8)), inference(subsumption_resolution, [], [f1978, f460])).
fof(f1978, plain, ((e1 = e2) | (~ spl4_7 | ~ spl4_8)), inference(backward_demodulation, [], [f643, f639])).
fof(f1977, plain, (~ spl4_8 | ~ spl4_10), inference(avatar_contradiction_clause, [], [f1976])).
fof(f1976, plain, ($false | (~ spl4_8 | ~ spl4_10)), inference(subsumption_resolution, [], [f1975, f464])).
fof(f1975, plain, ((e2 = e4) | (~ spl4_8 | ~ spl4_10)), inference(backward_demodulation, [], [f651, f643])).
fof(f1973, plain, (spl4_5 | ~ spl4_10 | ~ spl4_30), inference(avatar_split_clause, [], [f1970, f733, f649, f628])).
fof(f1970, plain, ((e4 = op(e4, e4)) | (~ spl4_10 | ~ spl4_30)), inference(backward_demodulation, [], [f1910, f651])).
fof(f1910, plain, ((e4 = op(e4, op(e4, e3))) | ~ spl4_30), inference(backward_demodulation, [], [f354, f735])).
fof(f1967, plain, (~ spl4_12 | ~ spl4_14), inference(avatar_contradiction_clause, [], [f1966])).
fof(f1966, plain, ($false | (~ spl4_12 | ~ spl4_14)), inference(subsumption_resolution, [], [f1965, f461])).
fof(f1965, plain, ((e1 = e3) | (~ spl4_12 | ~ spl4_14)), inference(backward_demodulation, [], [f668, f660])).
fof(f1963, plain, (~ spl4_14 | ~ spl4_15), inference(avatar_contradiction_clause, [], [f1962])).
fof(f1962, plain, ($false | (~ spl4_14 | ~ spl4_15)), inference(subsumption_resolution, [], [f1961, f465])).
fof(f1961, plain, ((e3 = e4) | (~ spl4_14 | ~ spl4_15)), inference(backward_demodulation, [], [f672, f668])).
fof(f1960, plain, (spl4_5 | ~ spl4_15 | ~ spl4_55), inference(avatar_split_clause, [], [f1956, f838, f670, f628])).
fof(f1956, plain, ((e4 = op(e4, e4)) | (~ spl4_15 | ~ spl4_55)), inference(backward_demodulation, [], [f1833, f672])).
fof(f1833, plain, ((e4 = op(e4, op(e4, e2))) | ~ spl4_55), inference(backward_demodulation, [], [f353, f840])).
fof(f1952, plain, (~ spl4_18 | ~ spl4_19), inference(avatar_contradiction_clause, [], [f1951])).
fof(f1951, plain, ($false | (~ spl4_18 | ~ spl4_19)), inference(subsumption_resolution, [], [f1950, f463])).
fof(f1950, plain, ((e2 = e3) | (~ spl4_18 | ~ spl4_19)), inference(backward_demodulation, [], [f689, f685])).
fof(f1949, plain, (~ spl4_19 | ~ spl4_20), inference(avatar_contradiction_clause, [], [f1948])).
fof(f1948, plain, ($false | (~ spl4_19 | ~ spl4_20)), inference(subsumption_resolution, [], [f1947, f465])).
fof(f1947, plain, ((e3 = e4) | (~ spl4_19 | ~ spl4_20)), inference(backward_demodulation, [], [f693, f689])).
fof(f1946, plain, (spl4_5 | ~ spl4_20 | ~ spl4_80), inference(avatar_split_clause, [], [f1941, f943, f691, f628])).
fof(f1941, plain, ((e4 = op(e4, e4)) | (~ spl4_20 | ~ spl4_80)), inference(backward_demodulation, [], [f1748, f693])).
fof(f1748, plain, ((e4 = op(e4, op(e4, e1))) | ~ spl4_80), inference(backward_demodulation, [], [f352, f945])).
fof(f1943, plain, (~ spl4_10 | ~ spl4_20), inference(avatar_split_clause, [], [f1938, f691, f649])).
fof(f1938, plain, (~ (e4 = op(e4, e3)) | ~ spl4_20), inference(backward_demodulation, [], [f451, f693])).
fof(f1936, plain, (~ spl4_22 | ~ spl4_23), inference(avatar_contradiction_clause, [], [f1935])).
fof(f1935, plain, ($false | (~ spl4_22 | ~ spl4_23)), inference(subsumption_resolution, [], [f1934, f460])).
fof(f1934, plain, ((e1 = e2) | (~ spl4_22 | ~ spl4_23)), inference(backward_demodulation, [], [f706, f702])).
fof(f1930, plain, (~ spl4_24 | ~ spl4_25), inference(avatar_contradiction_clause, [], [f1929])).
fof(f1929, plain, ($false | (~ spl4_24 | ~ spl4_25)), inference(subsumption_resolution, [], [f1928, f465])).
fof(f1928, plain, ((e3 = e4) | (~ spl4_24 | ~ spl4_25)), inference(backward_demodulation, [], [f714, f710])).
fof(f1927, plain, (spl4_5 | ~ spl4_25 | ~ spl4_105), inference(avatar_split_clause, [], [f1921, f1048, f712, f628])).
fof(f1921, plain, ((e4 = op(e4, e4)) | (~ spl4_25 | ~ spl4_105)), inference(backward_demodulation, [], [f1659, f714])).
fof(f1915, plain, (~ spl4_26 | ~ spl4_30), inference(avatar_contradiction_clause, [], [f1914])).
fof(f1914, plain, ($false | (~ spl4_26 | ~ spl4_30)), inference(subsumption_resolution, [], [f1913, f459])).
fof(f1913, plain, ((e0 = e4) | (~ spl4_26 | ~ spl4_30)), inference(backward_demodulation, [], [f735, f719])).
fof(f1900, plain, (~ spl4_29 | ~ spl4_34), inference(avatar_split_clause, [], [f1894, f750, f729])).
fof(f1894, plain, (~ (e3 = op(e3, e4)) | ~ spl4_34), inference(backward_demodulation, [], [f445, f752])).
fof(f1877, plain, (~ spl4_42 | ~ spl4_81 | ~ spl4_116), inference(avatar_contradiction_clause, [], [f1876])).
fof(f1876, plain, ($false | (~ spl4_42 | ~ spl4_81 | ~ spl4_116)), inference(subsumption_resolution, [], [f1875, f458])).
fof(f1875, plain, ((e0 = e3) | (~ spl4_42 | ~ spl4_81 | ~ spl4_116)), inference(forward_demodulation, [], [f1867, f1097])).
fof(f1867, plain, ((e3 = op(e0, e1)) | (~ spl4_42 | ~ spl4_81)), inference(backward_demodulation, [], [f1738, f786])).
fof(f1738, plain, ((e3 = op(e0, op(e3, e1))) | ~ spl4_81), inference(backward_demodulation, [], [f347, f950])).
fof(f1871, plain, (~ spl4_27 | ~ spl4_42), inference(avatar_split_clause, [], [f1865, f784, f721])).
fof(f1865, plain, (~ (e1 = op(e3, e4)) | ~ spl4_42), inference(backward_demodulation, [], [f442, f786])).
fof(f1846, plain, (~ spl4_52 | ~ spl4_53), inference(avatar_contradiction_clause, [], [f1845])).
fof(f1845, plain, ($false | (~ spl4_52 | ~ spl4_53)), inference(subsumption_resolution, [], [f1844, f460])).
fof(f1844, plain, ((e1 = e2) | (~ spl4_52 | ~ spl4_53)), inference(backward_demodulation, [], [f832, f828])).
fof(f1843, plain, (~ spl4_53 | ~ spl4_54), inference(avatar_contradiction_clause, [], [f1842])).
fof(f1842, plain, ($false | (~ spl4_53 | ~ spl4_54)), inference(subsumption_resolution, [], [f1841, f463])).
fof(f1841, plain, ((e2 = e3) | (~ spl4_53 | ~ spl4_54)), inference(backward_demodulation, [], [f836, f832])).
fof(f1840, plain, (~ spl4_54 | ~ spl4_55), inference(avatar_contradiction_clause, [], [f1839])).
fof(f1839, plain, ($false | (~ spl4_54 | ~ spl4_55)), inference(subsumption_resolution, [], [f1838, f465])).
fof(f1838, plain, ((e3 = e4) | (~ spl4_54 | ~ spl4_55)), inference(backward_demodulation, [], [f840, f836])).
fof(f1830, plain, (~ spl4_6 | ~ spl4_56), inference(avatar_split_clause, [], [f1827, f843, f633])).
fof(f1827, plain, (~ (e0 = op(e4, e3)) | ~ spl4_56), inference(backward_demodulation, [], [f394, f845])).
fof(f1810, plain, (~ spl4_61 | ~ spl4_121), inference(avatar_contradiction_clause, [], [f1809])).
fof(f1809, plain, ($false | (~ spl4_61 | ~ spl4_121)), inference(subsumption_resolution, [], [f1808, f457])).
fof(f1808, plain, ((e0 = e2) | (~ spl4_61 | ~ spl4_121)), inference(forward_demodulation, [], [f1799, f1118])).
fof(f1799, plain, ((op(e0, e0) = e2) | ~ spl4_61), inference(backward_demodulation, [], [f343, f866])).
fof(f1763, plain, (~ spl4_77 | ~ spl4_78), inference(avatar_contradiction_clause, [], [f1762])).
fof(f1762, plain, ($false | (~ spl4_77 | ~ spl4_78)), inference(subsumption_resolution, [], [f1761, f460])).
fof(f1761, plain, ((e1 = e2) | (~ spl4_77 | ~ spl4_78)), inference(backward_demodulation, [], [f937, f933])).
fof(f1760, plain, (~ spl4_78 | ~ spl4_79), inference(avatar_contradiction_clause, [], [f1759])).
fof(f1759, plain, ($false | (~ spl4_78 | ~ spl4_79)), inference(subsumption_resolution, [], [f1758, f463])).
fof(f1758, plain, ((e2 = e3) | (~ spl4_78 | ~ spl4_79)), inference(backward_demodulation, [], [f941, f937])).
fof(f1757, plain, (~ spl4_79 | ~ spl4_80), inference(avatar_contradiction_clause, [], [f1756])).
fof(f1756, plain, ($false | (~ spl4_79 | ~ spl4_80)), inference(subsumption_resolution, [], [f1755, f465])).
fof(f1755, plain, ((e3 = e4) | (~ spl4_79 | ~ spl4_80)), inference(backward_demodulation, [], [f945, f941])).
fof(f1714, plain, (~ spl4_81 | ~ spl4_91), inference(avatar_split_clause, [], [f1701, f990, f948])).
fof(f1701, plain, (~ (e0 = op(e1, e3)) | ~ spl4_91), inference(backward_demodulation, [], [f421, f992])).
fof(f1692, plain, (~ spl4_76 | ~ spl4_96), inference(avatar_split_clause, [], [f1683, f1011, f927])).
fof(f1683, plain, (~ (e0 = op(e1, e4)) | ~ spl4_96), inference(backward_demodulation, [], [f419, f1013])).
fof(f1676, plain, (~ spl4_102 | ~ spl4_103), inference(avatar_contradiction_clause, [], [f1675])).
fof(f1675, plain, ($false | (~ spl4_102 | ~ spl4_103)), inference(subsumption_resolution, [], [f1674, f460])).
fof(f1674, plain, ((e1 = e2) | (~ spl4_102 | ~ spl4_103)), inference(backward_demodulation, [], [f1042, f1038])).
fof(f1670, plain, (~ spl4_104 | ~ spl4_105), inference(avatar_contradiction_clause, [], [f1669])).
fof(f1669, plain, ($false | (~ spl4_104 | ~ spl4_105)), inference(subsumption_resolution, [], [f1668, f465])).
fof(f1668, plain, ((e3 = e4) | (~ spl4_104 | ~ spl4_105)), inference(backward_demodulation, [], [f1050, f1046])).
fof(f1653, plain, (~ spl4_81 | ~ spl4_106), inference(avatar_split_clause, [], [f1648, f1053, f948])).
fof(f1648, plain, (~ (e0 = op(e1, e3)) | ~ spl4_106), inference(backward_demodulation, [], [f386, f1055])).
fof(f1643, plain, (~ spl4_11 | ~ spl4_111), inference(avatar_split_clause, [], [f1637, f1074, f654])).
fof(f1637, plain, (~ (e0 = op(e4, e2)) | ~ spl4_111), inference(backward_demodulation, [], [f379, f1076])).
fof(f1629, plain, (~ spl4_111 | ~ spl4_116), inference(avatar_split_clause, [], [f1622, f1095, f1074])).
fof(f1622, plain, (~ (e0 = op(e0, e2)) | ~ spl4_116), inference(backward_demodulation, [], [f410, f1097])).
fof(f1607, plain, (~ spl4_101 | ~ spl4_121), inference(avatar_split_clause, [], [f1595, f1116, f1032])).
fof(f1595, plain, (~ (e0 = op(e0, e4)) | ~ spl4_121), inference(backward_demodulation, [], [f409, f1118])).
fof(f1604, plain, (~ spl4_116 | ~ spl4_121), inference(avatar_split_clause, [], [f1592, f1116, f1095])).
fof(f1592, plain, (~ (e0 = op(e0, e1)) | ~ spl4_121), inference(backward_demodulation, [], [f406, f1118])).
fof(f1546, plain, (~ spl4_184 | ~ spl4_157 | ~ spl4_93 | ~ spl4_181), inference(avatar_split_clause, [], [f586, f1497, f998, f1356, f1518])).
fof(f586, plain, (~ (e0 = op(e1, op(e1, e1))) | ~ (e2 = op(e1, e1)) | ~ (e4 = op(e1, op(e1, op(e1, e1)))) | ~ (e3 = op(op(e1, op(e1, e1)), op(e1, op(e1, e1))))), inference(cnf_transformation, [], [f222])).
fof(f222, plain, (~ (e0 = op(e1, op(e1, e1))) | ~ (e2 = op(e1, e1)) | ~ (e4 = op(e1, op(e1, op(e1, e1)))) | ~ (e3 = op(op(e1, op(e1, e1)), op(e1, op(e1, e1))))), inference(ennf_transformation, [], [f102])).
fof(f102, plain, ~ ((e0 = op(e1, op(e1, e1))) & (e2 = op(e1, e1)) & (e4 = op(e1, op(e1, op(e1, e1)))) & (e3 = op(op(e1, op(e1, e1)), op(e1, op(e1, e1))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG191+1.p', ax102)).
fof(f1541, plain, (~ spl4_184 | ~ spl4_157 | ~ spl4_91 | ~ spl4_153), inference(avatar_split_clause, [], [f581, f1336, f990, f1356, f1518])).
fof(f581, plain, (~ (e2 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, e1)) | ~ (e4 = op(e1, op(e1, op(e1, e1)))) | ~ (e3 = op(op(e1, op(e1, e1)), op(e1, op(e1, e1))))), inference(cnf_transformation, [], [f217])).
fof(f217, plain, (~ (e2 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, e1)) | ~ (e4 = op(e1, op(e1, op(e1, e1)))) | ~ (e3 = op(op(e1, op(e1, e1)), op(e1, op(e1, e1))))), inference(ennf_transformation, [], [f97])).
fof(f97, plain, ~ ((e2 = op(e1, op(e1, e1))) & (e0 = op(e1, e1)) & (e4 = op(e1, op(e1, op(e1, e1)))) & (e3 = op(op(e1, op(e1, e1)), op(e1, op(e1, e1))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG191+1.p', ax97)).
fof(f1539, plain, (~ spl4_185 | ~ spl4_165 | ~ spl4_125 | ~ spl4_180), inference(avatar_split_clause, [], [f579, f1492, f1132, f1400, f1528])).
fof(f579, plain, (~ (e1 = op(e0, op(e0, e0))) | ~ (op(e0, e0) = e4) | ~ (e2 = op(e0, op(e0, op(e0, e0)))) | ~ (e3 = op(op(e0, op(e0, e0)), op(e0, op(e0, e0))))), inference(cnf_transformation, [], [f215])).
fof(f215, plain, (~ (e1 = op(e0, op(e0, e0))) | ~ (op(e0, e0) = e4) | ~ (e2 = op(e0, op(e0, op(e0, e0)))) | ~ (e3 = op(op(e0, op(e0, e0)), op(e0, op(e0, e0))))), inference(ennf_transformation, [], [f95])).
fof(f95, plain, ~ ((e1 = op(e0, op(e0, e0))) & (op(e0, e0) = e4) & (e2 = op(e0, op(e0, op(e0, e0)))) & (e3 = op(op(e0, op(e0, e0)), op(e0, op(e0, e0))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG191+1.p', ax95)).
fof(f1538, plain, (~ spl4_183 | ~ spl4_143 | ~ spl4_2 | ~ spl4_167), inference(avatar_split_clause, [], [f578, f1409, f616, f1287, f1513])).
fof(f578, plain, (~ (e0 = op(e4, op(e4, e4))) | ~ (e1 = op(e4, e4)) | ~ (e2 = op(e4, op(e4, op(e4, e4)))) | ~ (e3 = op(op(e4, op(e4, e4)), op(e4, op(e4, e4))))), inference(cnf_transformation, [], [f214])).
fof(f214, plain, (~ (e0 = op(e4, op(e4, e4))) | ~ (e1 = op(e4, e4)) | ~ (e2 = op(e4, op(e4, op(e4, e4)))) | ~ (e3 = op(op(e4, op(e4, e4)), op(e4, op(e4, e4))))), inference(ennf_transformation, [], [f94])).
fof(f94, plain, ~ ((e0 = op(e4, op(e4, e4))) & (e1 = op(e4, e4)) & (e2 = op(e4, op(e4, op(e4, e4)))) & (e3 = op(op(e4, op(e4, e4)), op(e4, op(e4, e4))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG191+1.p', ax94)).
fof(f1535, plain, (~ spl4_184 | ~ spl4_145 | ~ spl4_91 | ~ spl4_146), inference(avatar_split_clause, [], [f575, f1300, f990, f1296, f1518])).
fof(f575, plain, (~ (e4 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, e1)) | ~ (e2 = op(e1, op(e1, op(e1, e1)))) | ~ (e3 = op(op(e1, op(e1, e1)), op(e1, op(e1, e1))))), inference(cnf_transformation, [], [f211])).
fof(f211, plain, (~ (e4 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, e1)) | ~ (e2 = op(e1, op(e1, op(e1, e1)))) | ~ (e3 = op(op(e1, op(e1, e1)), op(e1, op(e1, e1))))), inference(ennf_transformation, [], [f91])).
fof(f91, plain, ~ ((e4 = op(e1, op(e1, e1))) & (e0 = op(e1, e1)) & (e2 = op(e1, op(e1, op(e1, e1)))) & (e3 = op(op(e1, op(e1, e1)), op(e1, op(e1, e1))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG191+1.p', ax91)).
fof(f1526, plain, (~ spl4_183 | ~ spl4_134 | ~ spl4_1 | ~ spl4_139), inference(avatar_split_clause, [], [f570, f1267, f612, f1245, f1513])).
fof(f570, plain, (~ (e2 = op(e4, op(e4, e4))) | ~ (e0 = op(e4, e4)) | ~ (e1 = op(e4, op(e4, op(e4, e4)))) | ~ (e3 = op(op(e4, op(e4, e4)), op(e4, op(e4, e4))))), inference(cnf_transformation, [], [f206])).
fof(f206, plain, (~ (e2 = op(e4, op(e4, e4))) | ~ (e0 = op(e4, e4)) | ~ (e1 = op(e4, op(e4, op(e4, e4)))) | ~ (e3 = op(op(e4, op(e4, e4)), op(e4, op(e4, e4))))), inference(ennf_transformation, [], [f86])).
fof(f86, plain, ~ ((e2 = op(e4, op(e4, e4))) & (e0 = op(e4, e4)) & (e1 = op(e4, op(e4, op(e4, e4)))) & (e3 = op(op(e4, op(e4, e4)), op(e4, op(e4, e4))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG191+1.p', ax86)).
fof(f1523, plain, (~ spl4_184 | ~ spl4_177 | ~ spl4_95 | ~ spl4_153), inference(avatar_split_clause, [], [f567, f1336, f1006, f1466, f1518])).
fof(f567, plain, (~ (e2 = op(e1, op(e1, e1))) | ~ (e4 = op(e1, e1)) | ~ (e0 = op(e1, op(e1, op(e1, e1)))) | ~ (e3 = op(op(e1, op(e1, e1)), op(e1, op(e1, e1))))), inference(cnf_transformation, [], [f203])).
fof(f203, plain, (~ (e2 = op(e1, op(e1, e1))) | ~ (e4 = op(e1, e1)) | ~ (e0 = op(e1, op(e1, op(e1, e1)))) | ~ (e3 = op(op(e1, op(e1, e1)), op(e1, op(e1, e1))))), inference(ennf_transformation, [], [f83])).
fof(f83, plain, ~ ((e2 = op(e1, op(e1, e1))) & (e4 = op(e1, e1)) & (e0 = op(e1, op(e1, op(e1, e1)))) & (e3 = op(op(e1, op(e1, e1)), op(e1, op(e1, e1))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG191+1.p', ax83)).
fof(f1522, plain, (~ spl4_183 | ~ spl4_161 | ~ spl4_3 | ~ spl4_147), inference(avatar_split_clause, [], [f566, f1305, f620, f1377, f1513])).
fof(f566, plain, (~ (e1 = op(e4, op(e4, e4))) | ~ (e2 = op(e4, e4)) | ~ (e0 = op(e4, op(e4, op(e4, e4)))) | ~ (e3 = op(op(e4, op(e4, e4)), op(e4, op(e4, e4))))), inference(cnf_transformation, [], [f202])).
fof(f202, plain, (~ (e1 = op(e4, op(e4, e4))) | ~ (e2 = op(e4, e4)) | ~ (e0 = op(e4, op(e4, op(e4, e4)))) | ~ (e3 = op(op(e4, op(e4, e4)), op(e4, op(e4, e4))))), inference(ennf_transformation, [], [f82])).
fof(f82, plain, ~ ((e1 = op(e4, op(e4, e4))) & (e2 = op(e4, e4)) & (e0 = op(e4, op(e4, op(e4, e4)))) & (e3 = op(op(e4, op(e4, e4)), op(e4, op(e4, e4))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG191+1.p', ax82)).
fof(f1516, plain, (~ spl4_183 | ~ spl4_161 | ~ spl4_2 | ~ spl4_139), inference(avatar_split_clause, [], [f564, f1267, f616, f1377, f1513])).
fof(f564, plain, (~ (e2 = op(e4, op(e4, e4))) | ~ (e1 = op(e4, e4)) | ~ (e0 = op(e4, op(e4, op(e4, e4)))) | ~ (e3 = op(op(e4, op(e4, e4)), op(e4, op(e4, e4))))), inference(cnf_transformation, [], [f200])).
fof(f200, plain, (~ (e2 = op(e4, op(e4, e4))) | ~ (e1 = op(e4, e4)) | ~ (e0 = op(e4, op(e4, op(e4, e4)))) | ~ (e3 = op(op(e4, op(e4, e4)), op(e4, op(e4, e4))))), inference(ennf_transformation, [], [f80])).
fof(f80, plain, ~ ((e2 = op(e4, op(e4, e4))) & (e1 = op(e4, e4)) & (e0 = op(e4, op(e4, op(e4, e4)))) & (e3 = op(op(e4, op(e4, e4)), op(e4, op(e4, e4))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG191+1.p', ax80)).
fof(f1511, plain, (~ spl4_182 | ~ spl4_163 | ~ spl4_62 | ~ spl4_138), inference(avatar_split_clause, [], [f563, f1262, f868, f1386, f1508])).
fof(f563, plain, (~ (e4 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, e2)) | ~ (e0 = op(e2, op(e2, op(e2, e2)))) | ~ (e3 = op(op(e2, op(e2, e2)), op(e2, op(e2, e2))))), inference(cnf_transformation, [], [f199])).
fof(f199, plain, (~ (e4 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, e2)) | ~ (e0 = op(e2, op(e2, op(e2, e2)))) | ~ (e3 = op(op(e2, op(e2, e2)), op(e2, op(e2, e2))))), inference(ennf_transformation, [], [f79])).
fof(f79, plain, ~ ((e4 = op(e2, op(e2, e2))) & (e1 = op(e2, e2)) & (e0 = op(e2, op(e2, op(e2, e2)))) & (e3 = op(op(e2, op(e2, e2)), op(e2, op(e2, e2))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG191+1.p', ax79)).
fof(f1503, plain, (~ spl4_178 | ~ spl4_173 | ~ spl4_122 | ~ spl4_168), inference(avatar_split_clause, [], [f559, f1414, f1120, f1444, f1476])).
fof(f559, plain, (~ (e3 = op(e0, op(e0, e0))) | ~ (op(e0, e0) = e1) | ~ (e4 = op(e0, op(e0, op(e0, e0)))) | ~ (e2 = op(op(e0, op(e0, e0)), op(e0, op(e0, e0))))), inference(cnf_transformation, [], [f195])).
fof(f195, plain, (~ (e3 = op(e0, op(e0, e0))) | ~ (op(e0, e0) = e1) | ~ (e4 = op(e0, op(e0, op(e0, e0)))) | ~ (e2 = op(op(e0, op(e0, e0)), op(e0, op(e0, e0))))), inference(ennf_transformation, [], [f75])).
fof(f75, plain, ~ ((e3 = op(e0, op(e0, e0))) & (op(e0, e0) = e1) & (e4 = op(e0, op(e0, op(e0, e0)))) & (e2 = op(op(e0, op(e0, e0)), op(e0, op(e0, e0))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG191+1.p', ax75)).
fof(f1501, plain, (~ spl4_176 | ~ spl4_157 | ~ spl4_91 | ~ spl4_148), inference(avatar_split_clause, [], [f557, f1310, f990, f1356, f1462])).
fof(f557, plain, (~ (e3 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, e1)) | ~ (e4 = op(e1, op(e1, op(e1, e1)))) | ~ (e2 = op(op(e1, op(e1, e1)), op(e1, op(e1, e1))))), inference(cnf_transformation, [], [f193])).
fof(f193, plain, (~ (e3 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, e1)) | ~ (e4 = op(e1, op(e1, op(e1, e1)))) | ~ (e2 = op(op(e1, op(e1, e1)), op(e1, op(e1, e1))))), inference(ennf_transformation, [], [f73])).
fof(f73, plain, ~ ((e3 = op(e1, op(e1, e1))) & (e0 = op(e1, e1)) & (e4 = op(e1, op(e1, op(e1, e1)))) & (e2 = op(op(e1, op(e1, e1)), op(e1, op(e1, e1))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG191+1.p', ax73)).
fof(f1495, plain, (~ spl4_178 | ~ spl4_170 | ~ spl4_125 | ~ spl4_180), inference(avatar_split_clause, [], [f555, f1492, f1132, f1426, f1476])).
fof(f555, plain, (~ (e1 = op(e0, op(e0, e0))) | ~ (op(e0, e0) = e4) | ~ (e3 = op(e0, op(e0, op(e0, e0)))) | ~ (e2 = op(op(e0, op(e0, e0)), op(e0, op(e0, e0))))), inference(cnf_transformation, [], [f191])).
fof(f191, plain, (~ (e1 = op(e0, op(e0, e0))) | ~ (op(e0, e0) = e4) | ~ (e3 = op(e0, op(e0, op(e0, e0)))) | ~ (e2 = op(op(e0, op(e0, e0)), op(e0, op(e0, e0))))), inference(ennf_transformation, [], [f71])).
fof(f71, plain, ~ ((e1 = op(e0, op(e0, e0))) & (op(e0, e0) = e4) & (e3 = op(e0, op(e0, op(e0, e0)))) & (e2 = op(op(e0, op(e0, e0)), op(e0, op(e0, e0))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG191+1.p', ax71)).
fof(f1490, plain, (~ spl4_175 | ~ spl4_151 | ~ spl4_2 | ~ spl4_167), inference(avatar_split_clause, [], [f554, f1409, f616, f1325, f1457])).
fof(f554, plain, (~ (e0 = op(e4, op(e4, e4))) | ~ (e1 = op(e4, e4)) | ~ (e3 = op(e4, op(e4, op(e4, e4)))) | ~ (e2 = op(op(e4, op(e4, e4)), op(e4, op(e4, e4))))), inference(cnf_transformation, [], [f190])).
fof(f190, plain, (~ (e0 = op(e4, op(e4, e4))) | ~ (e1 = op(e4, e4)) | ~ (e3 = op(e4, op(e4, op(e4, e4)))) | ~ (e2 = op(op(e4, op(e4, e4)), op(e4, op(e4, e4))))), inference(ennf_transformation, [], [f70])).
fof(f70, plain, ~ ((e0 = op(e4, op(e4, e4))) & (e1 = op(e4, e4)) & (e3 = op(e4, op(e4, op(e4, e4)))) & (e2 = op(op(e4, op(e4, e4)), op(e4, op(e4, e4))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG191+1.p', ax70)).
fof(f1489, plain, (~ spl4_178 | ~ spl4_170 | ~ spl4_122 | ~ spl4_166), inference(avatar_split_clause, [], [f553, f1404, f1120, f1426, f1476])).
fof(f553, plain, (~ (e4 = op(e0, op(e0, e0))) | ~ (op(e0, e0) = e1) | ~ (e3 = op(e0, op(e0, op(e0, e0)))) | ~ (e2 = op(op(e0, op(e0, e0)), op(e0, op(e0, e0))))), inference(cnf_transformation, [], [f189])).
fof(f189, plain, (~ (e4 = op(e0, op(e0, e0))) | ~ (op(e0, e0) = e1) | ~ (e3 = op(e0, op(e0, op(e0, e0)))) | ~ (e2 = op(op(e0, op(e0, e0)), op(e0, op(e0, e0))))), inference(ennf_transformation, [], [f69])).
fof(f69, plain, ~ ((e4 = op(e0, op(e0, e0))) & (op(e0, e0) = e1) & (e3 = op(e0, op(e0, op(e0, e0)))) & (e2 = op(op(e0, op(e0, e0)), op(e0, op(e0, e0))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG191+1.p', ax69)).
fof(f1488, plain, (~ spl4_175 | ~ spl4_151 | ~ spl4_1 | ~ spl4_147), inference(avatar_split_clause, [], [f552, f1305, f612, f1325, f1457])).
fof(f552, plain, (~ (e1 = op(e4, op(e4, e4))) | ~ (e0 = op(e4, e4)) | ~ (e3 = op(e4, op(e4, op(e4, e4)))) | ~ (e2 = op(op(e4, op(e4, e4)), op(e4, op(e4, e4))))), inference(cnf_transformation, [], [f188])).
fof(f188, plain, (~ (e1 = op(e4, op(e4, e4))) | ~ (e0 = op(e4, e4)) | ~ (e3 = op(e4, op(e4, op(e4, e4)))) | ~ (e2 = op(op(e4, op(e4, e4)), op(e4, op(e4, e4))))), inference(ennf_transformation, [], [f68])).
fof(f68, plain, ~ ((e1 = op(e4, op(e4, e4))) & (e0 = op(e4, e4)) & (e3 = op(e4, op(e4, op(e4, e4)))) & (e2 = op(op(e4, op(e4, e4)), op(e4, op(e4, e4))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG191+1.p', ax68)).
fof(f1487, plain, (~ spl4_176 | ~ spl4_152 | ~ spl4_91 | ~ spl4_146), inference(avatar_split_clause, [], [f551, f1300, f990, f1330, f1462])).
fof(f551, plain, (~ (e4 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, e1)) | ~ (e3 = op(e1, op(e1, op(e1, e1)))) | ~ (e2 = op(op(e1, op(e1, e1)), op(e1, op(e1, e1))))), inference(cnf_transformation, [], [f187])).
fof(f187, plain, (~ (e4 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, e1)) | ~ (e3 = op(e1, op(e1, op(e1, e1)))) | ~ (e2 = op(op(e1, op(e1, e1)), op(e1, op(e1, e1))))), inference(ennf_transformation, [], [f67])).
fof(f67, plain, ~ ((e4 = op(e1, op(e1, e1))) & (e0 = op(e1, e1)) & (e3 = op(e1, op(e1, op(e1, e1)))) & (e2 = op(op(e1, op(e1, e1)), op(e1, op(e1, e1))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG191+1.p', ax67)).
fof(f1484, plain, (~ spl4_175 | ~ spl4_134 | ~ spl4_4 | ~ spl4_167), inference(avatar_split_clause, [], [f548, f1409, f624, f1245, f1457])).
fof(f548, plain, (~ (e0 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, e4)) | ~ (e1 = op(e4, op(e4, op(e4, e4)))) | ~ (e2 = op(op(e4, op(e4, e4)), op(e4, op(e4, e4))))), inference(cnf_transformation, [], [f184])).
fof(f184, plain, (~ (e0 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, e4)) | ~ (e1 = op(e4, op(e4, op(e4, e4)))) | ~ (e2 = op(op(e4, op(e4, e4)), op(e4, op(e4, e4))))), inference(ennf_transformation, [], [f64])).
fof(f64, plain, ~ ((e0 = op(e4, op(e4, e4))) & (e3 = op(e4, e4)) & (e1 = op(e4, op(e4, op(e4, e4)))) & (e2 = op(op(e4, op(e4, e4)), op(e4, op(e4, e4))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG191+1.p', ax64)).
fof(f1474, plain, (~ spl4_175 | ~ spl4_134 | ~ spl4_1 | ~ spl4_135), inference(avatar_split_clause, [], [f546, f1249, f612, f1245, f1457])).
fof(f546, plain, (~ (e3 = op(e4, op(e4, e4))) | ~ (e0 = op(e4, e4)) | ~ (e1 = op(e4, op(e4, op(e4, e4)))) | ~ (e2 = op(op(e4, op(e4, e4)), op(e4, op(e4, e4))))), inference(cnf_transformation, [], [f182])).
fof(f182, plain, (~ (e3 = op(e4, op(e4, e4))) | ~ (e0 = op(e4, e4)) | ~ (e1 = op(e4, op(e4, op(e4, e4)))) | ~ (e2 = op(op(e4, op(e4, e4)), op(e4, op(e4, e4))))), inference(ennf_transformation, [], [f62])).
fof(f62, plain, ~ ((e3 = op(e4, op(e4, e4))) & (e0 = op(e4, e4)) & (e1 = op(e4, op(e4, op(e4, e4)))) & (e2 = op(op(e4, op(e4, e4)), op(e4, op(e4, e4))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG191+1.p', ax62)).
fof(f1473, plain, (~ spl4_174 | ~ spl4_131 | ~ spl4_31 | ~ spl4_132), inference(avatar_split_clause, [], [f545, f1236, f738, f1232, f1452])).
fof(f545, plain, (~ (e4 = op(e3, op(e3, e3))) | ~ (e0 = op(e3, e3)) | ~ (e1 = op(e3, op(e3, op(e3, e3)))) | ~ (e2 = op(op(e3, op(e3, e3)), op(e3, op(e3, e3))))), inference(cnf_transformation, [], [f181])).
fof(f181, plain, (~ (e4 = op(e3, op(e3, e3))) | ~ (e0 = op(e3, e3)) | ~ (e1 = op(e3, op(e3, op(e3, e3)))) | ~ (e2 = op(op(e3, op(e3, e3)), op(e3, op(e3, e3))))), inference(ennf_transformation, [], [f61])).
fof(f61, plain, ~ ((e4 = op(e3, op(e3, e3))) & (e0 = op(e3, e3)) & (e1 = op(e3, op(e3, op(e3, e3)))) & (e2 = op(op(e3, op(e3, e3)), op(e3, op(e3, e3))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG191+1.p', ax61)).
fof(f1472, plain, (~ spl4_174 | ~ spl4_159 | ~ spl4_35 | ~ spl4_149), inference(avatar_split_clause, [], [f544, f1315, f754, f1368, f1452])).
fof(f544, plain, (~ (e1 = op(e3, op(e3, e3))) | ~ (e4 = op(e3, e3)) | ~ (e0 = op(e3, op(e3, op(e3, e3)))) | ~ (e2 = op(op(e3, op(e3, e3)), op(e3, op(e3, e3))))), inference(cnf_transformation, [], [f180])).
fof(f180, plain, (~ (e1 = op(e3, op(e3, e3))) | ~ (e4 = op(e3, e3)) | ~ (e0 = op(e3, op(e3, op(e3, e3)))) | ~ (e2 = op(op(e3, op(e3, e3)), op(e3, op(e3, e3))))), inference(ennf_transformation, [], [f60])).
fof(f60, plain, ~ ((e1 = op(e3, op(e3, e3))) & (e4 = op(e3, e3)) & (e0 = op(e3, op(e3, op(e3, e3)))) & (e2 = op(op(e3, op(e3, e3)), op(e3, op(e3, e3))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG191+1.p', ax60)).
fof(f1469, plain, (~ spl4_176 | ~ spl4_177 | ~ spl4_94 | ~ spl4_146), inference(avatar_split_clause, [], [f541, f1300, f1002, f1466, f1462])).
fof(f541, plain, (~ (e4 = op(e1, op(e1, e1))) | ~ (e3 = op(e1, e1)) | ~ (e0 = op(e1, op(e1, op(e1, e1)))) | ~ (e2 = op(op(e1, op(e1, e1)), op(e1, op(e1, e1))))), inference(cnf_transformation, [], [f177])).
fof(f177, plain, (~ (e4 = op(e1, op(e1, e1))) | ~ (e3 = op(e1, e1)) | ~ (e0 = op(e1, op(e1, op(e1, e1)))) | ~ (e2 = op(op(e1, op(e1, e1)), op(e1, op(e1, e1))))), inference(ennf_transformation, [], [f57])).
fof(f57, plain, ~ ((e4 = op(e1, op(e1, e1))) & (e3 = op(e1, e1)) & (e0 = op(e1, op(e1, op(e1, e1)))) & (e2 = op(op(e1, op(e1, e1)), op(e1, op(e1, e1))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG191+1.p', ax57)).
fof(f1447, plain, (~ spl4_164 | ~ spl4_173 | ~ spl4_123 | ~ spl4_168), inference(avatar_split_clause, [], [f535, f1414, f1124, f1444, f1396])).
fof(f535, plain, (~ (e3 = op(e0, op(e0, e0))) | ~ (op(e0, e0) = e2) | ~ (e4 = op(e0, op(e0, op(e0, e0)))) | ~ (e1 = op(op(e0, op(e0, e0)), op(e0, op(e0, e0))))), inference(cnf_transformation, [], [f171])).
fof(f171, plain, (~ (e3 = op(e0, op(e0, e0))) | ~ (op(e0, e0) = e2) | ~ (e4 = op(e0, op(e0, op(e0, e0)))) | ~ (e1 = op(op(e0, op(e0, e0)), op(e0, op(e0, e0))))), inference(ennf_transformation, [], [f51])).
fof(f51, plain, ~ ((e3 = op(e0, op(e0, e0))) & (op(e0, e0) = e2) & (e4 = op(e0, op(e0, op(e0, e0)))) & (e1 = op(op(e0, op(e0, e0)), op(e0, op(e0, e0))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG191+1.p', ax51)).
fof(f1442, plain, (~ spl4_158 | ~ spl4_156 | ~ spl4_31 | ~ spl4_141), inference(avatar_split_clause, [], [f534, f1277, f738, f1351, f1364])).
fof(f534, plain, (~ (e2 = op(e3, op(e3, e3))) | ~ (e0 = op(e3, e3)) | ~ (e4 = op(e3, op(e3, op(e3, e3)))) | ~ (e1 = op(op(e3, op(e3, e3)), op(e3, op(e3, e3))))), inference(cnf_transformation, [], [f170])).
fof(f170, plain, (~ (e2 = op(e3, op(e3, e3))) | ~ (e0 = op(e3, e3)) | ~ (e4 = op(e3, op(e3, op(e3, e3)))) | ~ (e1 = op(op(e3, op(e3, e3)), op(e3, op(e3, e3))))), inference(ennf_transformation, [], [f50])).
fof(f50, plain, ~ ((e2 = op(e3, op(e3, e3))) & (e0 = op(e3, e3)) & (e4 = op(e3, op(e3, op(e3, e3)))) & (e1 = op(op(e3, op(e3, e3)), op(e3, op(e3, e3))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG191+1.p', ax50)).
fof(f1441, plain, (~ spl4_162 | ~ spl4_155 | ~ spl4_61 | ~ spl4_140), inference(avatar_split_clause, [], [f533, f1272, f864, f1346, f1382])).
fof(f533, plain, (~ (e3 = op(e2, op(e2, e2))) | ~ (e0 = op(e2, e2)) | ~ (e4 = op(e2, op(e2, op(e2, e2)))) | ~ (e1 = op(op(e2, op(e2, e2)), op(e2, op(e2, e2))))), inference(cnf_transformation, [], [f169])).
fof(f169, plain, (~ (e3 = op(e2, op(e2, e2))) | ~ (e0 = op(e2, e2)) | ~ (e4 = op(e2, op(e2, op(e2, e2)))) | ~ (e1 = op(op(e2, op(e2, e2)), op(e2, op(e2, e2))))), inference(ennf_transformation, [], [f49])).
fof(f49, plain, ~ ((e3 = op(e2, op(e2, e2))) & (e0 = op(e2, e2)) & (e4 = op(e2, op(e2, op(e2, e2)))) & (e1 = op(op(e2, op(e2, e2)), op(e2, op(e2, e2))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG191+1.p', ax49)).
fof(f1435, plain, (~ spl4_164 | ~ spl4_170 | ~ spl4_125 | ~ spl4_171), inference(avatar_split_clause, [], [f531, f1432, f1132, f1426, f1396])).
fof(f531, plain, (~ (e2 = op(e0, op(e0, e0))) | ~ (op(e0, e0) = e4) | ~ (e3 = op(e0, op(e0, op(e0, e0)))) | ~ (e1 = op(op(e0, op(e0, e0)), op(e0, op(e0, e0))))), inference(cnf_transformation, [], [f167])).
fof(f167, plain, (~ (e2 = op(e0, op(e0, e0))) | ~ (op(e0, e0) = e4) | ~ (e3 = op(e0, op(e0, op(e0, e0)))) | ~ (e1 = op(op(e0, op(e0, e0)), op(e0, op(e0, e0))))), inference(ennf_transformation, [], [f47])).
fof(f47, plain, ~ ((e2 = op(e0, op(e0, e0))) & (op(e0, e0) = e4) & (e3 = op(e0, op(e0, op(e0, e0)))) & (e1 = op(op(e0, op(e0, e0)), op(e0, op(e0, e0))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG191+1.p', ax47)).
fof(f1422, plain, (~ spl4_158 | ~ spl4_142 | ~ spl4_35 | ~ spl4_169), inference(avatar_split_clause, [], [f526, f1419, f754, f1282, f1364])).
fof(f526, plain, (~ (e0 = op(e3, op(e3, e3))) | ~ (e4 = op(e3, e3)) | ~ (e2 = op(e3, op(e3, op(e3, e3)))) | ~ (e1 = op(op(e3, op(e3, e3)), op(e3, op(e3, e3))))), inference(cnf_transformation, [], [f162])).
fof(f162, plain, (~ (e0 = op(e3, op(e3, e3))) | ~ (e4 = op(e3, e3)) | ~ (e2 = op(e3, op(e3, op(e3, e3)))) | ~ (e1 = op(op(e3, op(e3, e3)), op(e3, op(e3, e3))))), inference(ennf_transformation, [], [f42])).
fof(f42, plain, ~ ((e0 = op(e3, op(e3, e3))) & (e4 = op(e3, e3)) & (e2 = op(e3, op(e3, op(e3, e3)))) & (e1 = op(op(e3, op(e3, e3)), op(e3, op(e3, e3))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG191+1.p', ax42)).
fof(f1417, plain, (~ spl4_164 | ~ spl4_165 | ~ spl4_125 | ~ spl4_168), inference(avatar_split_clause, [], [f525, f1414, f1132, f1400, f1396])).
fof(f525, plain, (~ (e3 = op(e0, op(e0, e0))) | ~ (op(e0, e0) = e4) | ~ (e2 = op(e0, op(e0, op(e0, e0)))) | ~ (e1 = op(op(e0, op(e0, e0)), op(e0, op(e0, e0))))), inference(cnf_transformation, [], [f161])).
fof(f161, plain, (~ (e3 = op(e0, op(e0, e0))) | ~ (op(e0, e0) = e4) | ~ (e2 = op(e0, op(e0, op(e0, e0)))) | ~ (e1 = op(op(e0, op(e0, e0)), op(e0, op(e0, e0))))), inference(ennf_transformation, [], [f41])).
fof(f41, plain, ~ ((e3 = op(e0, op(e0, e0))) & (op(e0, e0) = e4) & (e2 = op(e0, op(e0, op(e0, e0)))) & (e1 = op(op(e0, op(e0, e0)), op(e0, op(e0, e0))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG191+1.p', ax41)).
fof(f1394, plain, (~ spl4_160 | ~ spl4_143 | ~ spl4_1 | ~ spl4_135), inference(avatar_split_clause, [], [f522, f1249, f612, f1287, f1373])).
fof(f522, plain, (~ (e3 = op(e4, op(e4, e4))) | ~ (e0 = op(e4, e4)) | ~ (e2 = op(e4, op(e4, op(e4, e4)))) | ~ (e1 = op(op(e4, op(e4, e4)), op(e4, op(e4, e4))))), inference(cnf_transformation, [], [f158])).
fof(f158, plain, (~ (e3 = op(e4, op(e4, e4))) | ~ (e0 = op(e4, e4)) | ~ (e2 = op(e4, op(e4, op(e4, e4)))) | ~ (e1 = op(op(e4, op(e4, e4)), op(e4, op(e4, e4))))), inference(ennf_transformation, [], [f38])).
fof(f38, plain, ~ ((e3 = op(e4, op(e4, e4))) & (e0 = op(e4, e4)) & (e2 = op(e4, op(e4, op(e4, e4)))) & (e1 = op(op(e4, op(e4, e4)), op(e4, op(e4, e4))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG191+1.p', ax38)).
fof(f1371, plain, (~ spl4_158 | ~ spl4_159 | ~ spl4_33 | ~ spl4_132), inference(avatar_split_clause, [], [f515, f1236, f746, f1368, f1364])).
fof(f515, plain, (~ (e4 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, e3)) | ~ (e0 = op(e3, op(e3, op(e3, e3)))) | ~ (e1 = op(op(e3, op(e3, e3)), op(e3, op(e3, e3))))), inference(cnf_transformation, [], [f151])).
fof(f151, plain, (~ (e4 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, e3)) | ~ (e0 = op(e3, op(e3, op(e3, e3)))) | ~ (e1 = op(op(e3, op(e3, e3)), op(e3, op(e3, e3))))), inference(ennf_transformation, [], [f31])).
fof(f31, plain, ~ ((e4 = op(e3, op(e3, e3))) & (e2 = op(e3, e3)) & (e0 = op(e3, op(e3, op(e3, e3)))) & (e1 = op(op(e3, op(e3, e3)), op(e3, op(e3, e3))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG191+1.p', ax31)).
fof(f1362, plain, (~ spl4_136 | ~ spl4_155 | ~ spl4_64 | ~ spl4_154), inference(avatar_split_clause, [], [f514, f1341, f876, f1346, f1254])).
fof(f514, plain, (~ (e1 = op(e2, op(e2, e2))) | ~ (e3 = op(e2, e2)) | ~ (e4 = op(e2, op(e2, op(e2, e2)))) | ~ (e0 = op(op(e2, op(e2, e2)), op(e2, op(e2, e2))))), inference(cnf_transformation, [], [f150])).
fof(f150, plain, (~ (e1 = op(e2, op(e2, e2))) | ~ (e3 = op(e2, e2)) | ~ (e4 = op(e2, op(e2, op(e2, e2)))) | ~ (e0 = op(op(e2, op(e2, e2)), op(e2, op(e2, e2))))), inference(ennf_transformation, [], [f30])).
fof(f30, plain, ~ ((e1 = op(e2, op(e2, e2))) & (e3 = op(e2, e2)) & (e4 = op(e2, op(e2, op(e2, e2)))) & (e0 = op(op(e2, op(e2, e2)), op(e2, op(e2, e2))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG191+1.p', ax30)).
fof(f1361, plain, (~ spl4_144 | ~ spl4_157 | ~ spl4_94 | ~ spl4_153), inference(avatar_split_clause, [], [f513, f1336, f1002, f1356, f1292])).
fof(f513, plain, (~ (e2 = op(e1, op(e1, e1))) | ~ (e3 = op(e1, e1)) | ~ (e4 = op(e1, op(e1, op(e1, e1)))) | ~ (e0 = op(op(e1, op(e1, e1)), op(e1, op(e1, e1))))), inference(cnf_transformation, [], [f149])).
fof(f149, plain, (~ (e2 = op(e1, op(e1, e1))) | ~ (e3 = op(e1, e1)) | ~ (e4 = op(e1, op(e1, op(e1, e1)))) | ~ (e0 = op(op(e1, op(e1, e1)), op(e1, op(e1, e1))))), inference(ennf_transformation, [], [f29])).
fof(f29, plain, ~ ((e2 = op(e1, op(e1, e1))) & (e3 = op(e1, e1)) & (e4 = op(e1, op(e1, op(e1, e1)))) & (e0 = op(op(e1, op(e1, e1)), op(e1, op(e1, e1))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG191+1.p', ax29)).
fof(f1328, plain, (~ spl4_133 | ~ spl4_151 | ~ spl4_2 | ~ spl4_139), inference(avatar_split_clause, [], [f504, f1267, f616, f1325, f1241])).
fof(f504, plain, (~ (e2 = op(e4, op(e4, e4))) | ~ (e1 = op(e4, e4)) | ~ (e3 = op(e4, op(e4, op(e4, e4)))) | ~ (e0 = op(op(e4, op(e4, e4)), op(e4, op(e4, e4))))), inference(cnf_transformation, [], [f140])).
fof(f140, plain, (~ (e2 = op(e4, op(e4, e4))) | ~ (e1 = op(e4, e4)) | ~ (e3 = op(e4, op(e4, op(e4, e4)))) | ~ (e0 = op(op(e4, op(e4, e4)), op(e4, op(e4, e4))))), inference(ennf_transformation, [], [f20])).
fof(f20, plain, ~ ((e2 = op(e4, op(e4, e4))) & (e1 = op(e4, e4)) & (e3 = op(e4, op(e4, op(e4, e4)))) & (e0 = op(op(e4, op(e4, e4)), op(e4, op(e4, e4))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG191+1.p', ax20)).
fof(f1313, plain, (~ spl4_144 | ~ spl4_145 | ~ spl4_95 | ~ spl4_148), inference(avatar_split_clause, [], [f501, f1310, f1006, f1296, f1292])).
fof(f501, plain, (~ (e3 = op(e1, op(e1, e1))) | ~ (e4 = op(e1, e1)) | ~ (e2 = op(e1, op(e1, op(e1, e1)))) | ~ (e0 = op(op(e1, op(e1, e1)), op(e1, op(e1, e1))))), inference(cnf_transformation, [], [f137])).
fof(f137, plain, (~ (e3 = op(e1, op(e1, e1))) | ~ (e4 = op(e1, e1)) | ~ (e2 = op(e1, op(e1, op(e1, e1)))) | ~ (e0 = op(op(e1, op(e1, e1)), op(e1, op(e1, e1))))), inference(ennf_transformation, [], [f17])).
fof(f17, plain, ~ ((e3 = op(e1, op(e1, e1))) & (e4 = op(e1, e1)) & (e2 = op(e1, op(e1, op(e1, e1)))) & (e0 = op(op(e1, op(e1, e1)), op(e1, op(e1, e1))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG191+1.p', ax17)).
fof(f1290, plain, (~ spl4_133 | ~ spl4_143 | ~ spl4_2 | ~ spl4_135), inference(avatar_split_clause, [], [f498, f1249, f616, f1287, f1241])).
fof(f498, plain, (~ (e3 = op(e4, op(e4, e4))) | ~ (e1 = op(e4, e4)) | ~ (e2 = op(e4, op(e4, op(e4, e4)))) | ~ (e0 = op(op(e4, op(e4, e4)), op(e4, op(e4, e4))))), inference(cnf_transformation, [], [f134])).
fof(f134, plain, (~ (e3 = op(e4, op(e4, e4))) | ~ (e1 = op(e4, e4)) | ~ (e2 = op(e4, op(e4, op(e4, e4)))) | ~ (e0 = op(op(e4, op(e4, e4)), op(e4, op(e4, e4))))), inference(ennf_transformation, [], [f14])).
fof(f14, plain, ~ ((e3 = op(e4, op(e4, e4))) & (e1 = op(e4, e4)) & (e2 = op(e4, op(e4, op(e4, e4)))) & (e0 = op(op(e4, op(e4, e4)), op(e4, op(e4, e4))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG191+1.p', ax14)).
fof(f1280, plain, (~ spl4_130 | ~ spl4_131 | ~ spl4_35 | ~ spl4_141), inference(avatar_split_clause, [], [f496, f1277, f754, f1232, f1228])).
fof(f496, plain, (~ (e2 = op(e3, op(e3, e3))) | ~ (e4 = op(e3, e3)) | ~ (e1 = op(e3, op(e3, op(e3, e3)))) | ~ (e0 = op(op(e3, op(e3, e3)), op(e3, op(e3, e3))))), inference(cnf_transformation, [], [f132])).
fof(f132, plain, (~ (e2 = op(e3, op(e3, e3))) | ~ (e4 = op(e3, e3)) | ~ (e1 = op(e3, op(e3, op(e3, e3)))) | ~ (e0 = op(op(e3, op(e3, e3)), op(e3, op(e3, e3))))), inference(ennf_transformation, [], [f12])).
fof(f12, plain, ~ ((e2 = op(e3, op(e3, e3))) & (e4 = op(e3, e3)) & (e1 = op(e3, op(e3, op(e3, e3)))) & (e0 = op(op(e3, op(e3, e3)), op(e3, op(e3, e3))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG191+1.p', ax12)).
fof(f1270, plain, (~ spl4_133 | ~ spl4_134 | ~ spl4_4 | ~ spl4_139), inference(avatar_split_clause, [], [f494, f1267, f624, f1245, f1241])).
fof(f494, plain, (~ (e2 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, e4)) | ~ (e1 = op(e4, op(e4, op(e4, e4)))) | ~ (e0 = op(op(e4, op(e4, e4)), op(e4, op(e4, e4))))), inference(cnf_transformation, [], [f130])).
fof(f130, plain, (~ (e2 = op(e4, op(e4, e4))) | ~ (e3 = op(e4, e4)) | ~ (e1 = op(e4, op(e4, op(e4, e4)))) | ~ (e0 = op(op(e4, op(e4, e4)), op(e4, op(e4, e4))))), inference(ennf_transformation, [], [f10])).
fof(f10, plain, ~ ((e2 = op(e4, op(e4, e4))) & (e3 = op(e4, e4)) & (e1 = op(e4, op(e4, op(e4, e4)))) & (e0 = op(op(e4, op(e4, e4)), op(e4, op(e4, e4))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG191+1.p', ax10)).
fof(f1252, plain, (~ spl4_133 | ~ spl4_134 | ~ spl4_3 | ~ spl4_135), inference(avatar_split_clause, [], [f492, f1249, f620, f1245, f1241])).
fof(f492, plain, (~ (e3 = op(e4, op(e4, e4))) | ~ (e2 = op(e4, e4)) | ~ (e1 = op(e4, op(e4, op(e4, e4)))) | ~ (e0 = op(op(e4, op(e4, e4)), op(e4, op(e4, e4))))), inference(cnf_transformation, [], [f128])).
fof(f128, plain, (~ (e3 = op(e4, op(e4, e4))) | ~ (e2 = op(e4, e4)) | ~ (e1 = op(e4, op(e4, op(e4, e4)))) | ~ (e0 = op(op(e4, op(e4, e4)), op(e4, op(e4, e4))))), inference(ennf_transformation, [], [f8])).
fof(f8, plain, ~ ((e3 = op(e4, op(e4, e4))) & (e2 = op(e4, e4)) & (e1 = op(e4, op(e4, op(e4, e4)))) & (e0 = op(op(e4, op(e4, e4)), op(e4, op(e4, e4))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG191+1.p', ax8)).
fof(f1226, plain, (spl4_129 | spl4_128 | spl4_127 | spl4_126 | ~ spl4_21), inference(avatar_split_clause, [], [f486, f696, f1187, f1196, f1205, f1214])).
fof(f1214, plain, (spl4_129 <=> sP0), introduced(avatar_definition, [new_symbols(naming, [spl4_129])])).
fof(f1205, plain, (spl4_128 <=> sP1), introduced(avatar_definition, [new_symbols(naming, [spl4_128])])).
fof(f1196, plain, (spl4_127 <=> sP2), introduced(avatar_definition, [new_symbols(naming, [spl4_127])])).
fof(f1187, plain, (spl4_126 <=> sP3), introduced(avatar_definition, [new_symbols(naming, [spl4_126])])).
fof(f486, plain, (~ (e0 = op(e4, e0)) | sP3 | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f251])).
fof(f251, plain, ((~ (e4 = op(e4, e4)) & ~ (e3 = op(e4, e3)) & ~ (e2 = op(e4, e2)) & ~ (e1 = op(e4, e1)) & ~ (e0 = op(e4, e0))) | sP3 | sP2 | sP1 | sP0), inference(definition_folding, [], [f6, e250, e249, e248, e247])).
fof(f247, plain, ((~ (e4 = op(e0, e4)) & ~ (e3 = op(e0, e3)) & ~ (e2 = op(e0, e2)) & ~ (e1 = op(e0, e1)) & ~ (e0 = op(e0, e0))) | ~ sP0), inference(usedef, [], [e247])).
fof(e247, plain, (sP0 <=> (~ (e4 = op(e0, e4)) & ~ (e3 = op(e0, e3)) & ~ (e2 = op(e0, e2)) & ~ (e1 = op(e0, e1)) & ~ (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f248, plain, ((~ (e4 = op(e1, e4)) & ~ (e3 = op(e1, e3)) & ~ (e2 = op(e1, e2)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e1, e0))) | ~ sP1), inference(usedef, [], [e248])).
fof(e248, plain, (sP1 <=> (~ (e4 = op(e1, e4)) & ~ (e3 = op(e1, e3)) & ~ (e2 = op(e1, e2)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f249, plain, ((~ (e4 = op(e2, e4)) & ~ (e3 = op(e2, e3)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e2, e1)) & ~ (e0 = op(e2, e0))) | ~ sP2), inference(usedef, [], [e249])).
fof(e249, plain, (sP2 <=> (~ (e4 = op(e2, e4)) & ~ (e3 = op(e2, e3)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e2, e1)) & ~ (e0 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f250, plain, ((~ (e4 = op(e3, e4)) & ~ (e3 = op(e3, e3)) & ~ (e2 = op(e3, e2)) & ~ (e1 = op(e3, e1)) & ~ (e0 = op(e3, e0))) | ~ sP3), inference(usedef, [], [e250])).
fof(e250, plain, (sP3 <=> (~ (e4 = op(e3, e4)) & ~ (e3 = op(e3, e3)) & ~ (e2 = op(e3, e2)) & ~ (e1 = op(e3, e1)) & ~ (e0 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f6, plain, ((~ (e4 = op(e4, e4)) & ~ (e3 = op(e4, e3)) & ~ (e2 = op(e4, e2)) & ~ (e1 = op(e4, e1)) & ~ (e0 = op(e4, e0))) | (~ (e4 = op(e3, e4)) & ~ (e3 = op(e3, e3)) & ~ (e2 = op(e3, e2)) & ~ (e1 = op(e3, e1)) & ~ (e0 = op(e3, e0))) | (~ (e4 = op(e2, e4)) & ~ (e3 = op(e2, e3)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e2, e1)) & ~ (e0 = op(e2, e0))) | (~ (e4 = op(e1, e4)) & ~ (e3 = op(e1, e3)) & ~ (e2 = op(e1, e2)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e1, e0))) | (~ (e4 = op(e0, e4)) & ~ (e3 = op(e0, e3)) & ~ (e2 = op(e0, e2)) & ~ (e1 = op(e0, e1)) & ~ (e0 = op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG191+1.p', ax6)).
fof(f1224, plain, (spl4_129 | spl4_128 | spl4_127 | spl4_126 | ~ spl4_13), inference(avatar_split_clause, [], [f488, f662, f1187, f1196, f1205, f1214])).
fof(f488, plain, (~ (e2 = op(e4, e2)) | sP3 | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f251])).
fof(f1223, plain, (spl4_129 | spl4_128 | spl4_127 | spl4_126 | ~ spl4_9), inference(avatar_split_clause, [], [f489, f645, f1187, f1196, f1205, f1214])).
fof(f489, plain, (~ (e3 = op(e4, e3)) | sP3 | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f251])).
fof(f1222, plain, (spl4_129 | spl4_128 | spl4_127 | spl4_126 | ~ spl4_5), inference(avatar_split_clause, [], [f490, f628, f1187, f1196, f1205, f1214])).
fof(f490, plain, (~ (e4 = op(e4, e4)) | sP3 | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f251])).
fof(f1221, plain, (~ spl4_129 | ~ spl4_121), inference(avatar_split_clause, [], [f481, f1116, f1214])).
fof(f481, plain, (~ (e0 = op(e0, e0)) | ~ sP0), inference(cnf_transformation, [], [f255])).
fof(f255, plain, ((~ (e4 = op(e0, e4)) & ~ (e3 = op(e0, e3)) & ~ (e2 = op(e0, e2)) & ~ (e1 = op(e0, e1)) & ~ (e0 = op(e0, e0))) | ~ sP0), inference(nnf_transformation, [], [f247])).
fof(f1220, plain, (~ spl4_129 | ~ spl4_117), inference(avatar_split_clause, [], [f482, f1099, f1214])).
fof(f482, plain, (~ (e1 = op(e0, e1)) | ~ sP0), inference(cnf_transformation, [], [f255])).
fof(f1217, plain, (~ spl4_129 | ~ spl4_105), inference(avatar_split_clause, [], [f485, f1048, f1214])).
fof(f485, plain, (~ (e4 = op(e0, e4)) | ~ sP0), inference(cnf_transformation, [], [f255])).
fof(f1212, plain, (~ spl4_128 | ~ spl4_96), inference(avatar_split_clause, [], [f476, f1011, f1205])).
fof(f476, plain, (~ (e0 = op(e1, e0)) | ~ sP1), inference(cnf_transformation, [], [f254])).
fof(f254, plain, ((~ (e4 = op(e1, e4)) & ~ (e3 = op(e1, e3)) & ~ (e2 = op(e1, e2)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e1, e0))) | ~ sP1), inference(nnf_transformation, [], [f248])).
fof(f1211, plain, (~ spl4_128 | ~ spl4_92), inference(avatar_split_clause, [], [f477, f994, f1205])).
fof(f477, plain, (~ (e1 = op(e1, e1)) | ~ sP1), inference(cnf_transformation, [], [f254])).
fof(f1210, plain, (~ spl4_128 | ~ spl4_88), inference(avatar_split_clause, [], [f478, f977, f1205])).
fof(f478, plain, (~ (e2 = op(e1, e2)) | ~ sP1), inference(cnf_transformation, [], [f254])).
fof(f1203, plain, (~ spl4_127 | ~ spl4_71), inference(avatar_split_clause, [], [f471, f906, f1196])).
fof(f471, plain, (~ (e0 = op(e2, e0)) | ~ sP2), inference(cnf_transformation, [], [f253])).
fof(f253, plain, ((~ (e4 = op(e2, e4)) & ~ (e3 = op(e2, e3)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e2, e1)) & ~ (e0 = op(e2, e0))) | ~ sP2), inference(nnf_transformation, [], [f249])).
fof(f1202, plain, (~ spl4_127 | ~ spl4_67), inference(avatar_split_clause, [], [f472, f889, f1196])).
fof(f472, plain, (~ (e1 = op(e2, e1)) | ~ sP2), inference(cnf_transformation, [], [f253])).
fof(f1201, plain, (~ spl4_127 | ~ spl4_63), inference(avatar_split_clause, [], [f473, f872, f1196])).
fof(f473, plain, (~ (e2 = op(e2, e2)) | ~ sP2), inference(cnf_transformation, [], [f253])).
fof(f1199, plain, (~ spl4_127 | ~ spl4_55), inference(avatar_split_clause, [], [f475, f838, f1196])).
fof(f475, plain, (~ (e4 = op(e2, e4)) | ~ sP2), inference(cnf_transformation, [], [f253])).
fof(f1192, plain, (~ spl4_126 | ~ spl4_38), inference(avatar_split_clause, [], [f468, f767, f1187])).
fof(f468, plain, (~ (e2 = op(e3, e2)) | ~ sP3), inference(cnf_transformation, [], [f252])).
fof(f252, plain, ((~ (e4 = op(e3, e4)) & ~ (e3 = op(e3, e3)) & ~ (e2 = op(e3, e2)) & ~ (e1 = op(e3, e1)) & ~ (e0 = op(e3, e0))) | ~ sP3), inference(nnf_transformation, [], [f250])).
fof(f1191, plain, (~ spl4_126 | ~ spl4_34), inference(avatar_split_clause, [], [f469, f750, f1187])).
fof(f469, plain, (~ (e3 = op(e3, e3)) | ~ sP3), inference(cnf_transformation, [], [f252])).
fof(f1185, plain, (spl4_121 | spl4_116 | spl4_111 | spl4_106 | spl4_101), inference(avatar_split_clause, [], [f281, f1032, f1053, f1074, f1095, f1116])).
fof(f281, plain, ((e0 = op(e0, e4)) | (e0 = op(e0, e3)) | (e0 = op(e0, e2)) | (e0 = op(e0, e1)) | (e0 = op(e0, e0))), inference(cnf_transformation, [], [f2])).
fof(f2, plain, (((e4 = op(e4, e4)) | (e4 = op(e3, e4)) | (e4 = op(e2, e4)) | (e4 = op(e1, e4)) | (e4 = op(e0, e4))) & ((e4 = op(e4, e4)) | (e4 = op(e4, e3)) | (e4 = op(e4, e2)) | (e4 = op(e4, e1)) | (e4 = op(e4, e0))) & ((e3 = op(e4, e4)) | (e3 = op(e3, e4)) | (e3 = op(e2, e4)) | (e3 = op(e1, e4)) | (e3 = op(e0, e4))) & ((e3 = op(e4, e4)) | (e3 = op(e4, e3)) | (e3 = op(e4, e2)) | (e3 = op(e4, e1)) | (e3 = op(e4, e0))) & ((e2 = op(e4, e4)) | (e2 = op(e3, e4)) | (e2 = op(e2, e4)) | (e2 = op(e1, e4)) | (e2 = op(e0, e4))) & ((e2 = op(e4, e4)) | (e2 = op(e4, e3)) | (e2 = op(e4, e2)) | (e2 = op(e4, e1)) | (e2 = op(e4, e0))) & ((e1 = op(e4, e4)) | (e1 = op(e3, e4)) | (e1 = op(e2, e4)) | (e1 = op(e1, e4)) | (e1 = op(e0, e4))) & ((e1 = op(e4, e4)) | (e1 = op(e4, e3)) | (e1 = op(e4, e2)) | (e1 = op(e4, e1)) | (e1 = op(e4, e0))) & ((e0 = op(e4, e4)) | (e0 = op(e3, e4)) | (e0 = op(e2, e4)) | (e0 = op(e1, e4)) | (e0 = op(e0, e4))) & ((e0 = op(e4, e4)) | (e0 = op(e4, e3)) | (e0 = op(e4, e2)) | (e0 = op(e4, e1)) | (e0 = op(e4, e0))) & ((e4 = op(e4, e3)) | (e4 = op(e3, e3)) | (e4 = op(e2, e3)) | (e4 = op(e1, e3)) | (e4 = op(e0, e3))) & ((e4 = op(e3, e4)) | (e4 = op(e3, e3)) | (e4 = op(e3, e2)) | (e4 = op(e3, e1)) | (e4 = op(e3, e0))) & ((e3 = op(e4, e3)) | (e3 = op(e3, e3)) | (e3 = op(e2, e3)) | (e3 = op(e1, e3)) | (e3 = op(e0, e3))) & ((e3 = op(e3, e4)) | (e3 = op(e3, e3)) | (e3 = op(e3, e2)) | (e3 = op(e3, e1)) | (e3 = op(e3, e0))) & ((e2 = op(e4, e3)) | (e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(e0, e3))) & ((e2 = op(e3, e4)) | (e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, e0))) & ((e1 = op(e4, e3)) | (e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(e0, e3))) & ((e1 = op(e3, e4)) | (e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, e0))) & ((e0 = op(e4, e3)) | (e0 = op(e3, e3)) | (e0 = op(e2, e3)) | (e0 = op(e1, e3)) | (e0 = op(e0, e3))) & ((e0 = op(e3, e4)) | (e0 = op(e3, e3)) | (e0 = op(e3, e2)) | (e0 = op(e3, e1)) | (e0 = op(e3, e0))) & ((e4 = op(e4, e2)) | (e4 = op(e3, e2)) | (e4 = op(e2, e2)) | (e4 = op(e1, e2)) | (e4 = op(e0, e2))) & ((e4 = op(e2, e4)) | (e4 = op(e2, e3)) | (e4 = op(e2, e2)) | (e4 = op(e2, e1)) | (e4 = op(e2, e0))) & ((e3 = op(e4, e2)) | (e3 = op(e3, e2)) | (e3 = op(e2, e2)) | (e3 = op(e1, e2)) | (e3 = op(e0, e2))) & ((e3 = op(e2, e4)) | (e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, e0))) & ((e2 = op(e4, e2)) | (e2 = op(e3, e2)) | (e2 = op(e2, e2)) | (e2 = op(e1, e2)) | (e2 = op(e0, e2))) & ((e2 = op(e2, e4)) | (e2 = op(e2, e3)) | (e2 = op(e2, e2)) | (e2 = op(e2, e1)) | (e2 = op(e2, e0))) & ((e1 = op(e4, e2)) | (e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(e0, e2))) & ((e1 = op(e2, e4)) | (e1 = op(e2, e3)) | (e1 = op(e2, e2)) | (e1 = op(e2, e1)) | (e1 = op(e2, e0))) & ((e0 = op(e4, e2)) | (e0 = op(e3, e2)) | (e0 = op(e2, e2)) | (e0 = op(e1, e2)) | (e0 = op(e0, e2))) & ((e0 = op(e2, e4)) | (e0 = op(e2, e3)) | (e0 = op(e2, e2)) | (e0 = op(e2, e1)) | (e0 = op(e2, e0))) & ((e4 = op(e4, e1)) | (e4 = op(e3, e1)) | (e4 = op(e2, e1)) | (e4 = op(e1, e1)) | (e4 = op(e0, e1))) & ((e4 = op(e1, e4)) | (e4 = op(e1, e3)) | (e4 = op(e1, e2)) | (e4 = op(e1, e1)) | (e4 = op(e1, e0))) & ((e3 = op(e4, e1)) | (e3 = op(e3, e1)) | (e3 = op(e2, e1)) | (e3 = op(e1, e1)) | (e3 = op(e0, e1))) & ((e3 = op(e1, e4)) | (e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, e0))) & ((e2 = op(e4, e1)) | (e2 = op(e3, e1)) | (e2 = op(e2, e1)) | (e2 = op(e1, e1)) | (e2 = op(e0, e1))) & ((e2 = op(e1, e4)) | (e2 = op(e1, e3)) | (e2 = op(e1, e2)) | (e2 = op(e1, e1)) | (e2 = op(e1, e0))) & ((e1 = op(e4, e1)) | (e1 = op(e3, e1)) | (e1 = op(e2, e1)) | (e1 = op(e1, e1)) | (e1 = op(e0, e1))) & ((e1 = op(e1, e4)) | (e1 = op(e1, e3)) | (e1 = op(e1, e2)) | (e1 = op(e1, e1)) | (e1 = op(e1, e0))) & ((e0 = op(e4, e1)) | (e0 = op(e3, e1)) | (e0 = op(e2, e1)) | (e0 = op(e1, e1)) | (e0 = op(e0, e1))) & ((e0 = op(e1, e4)) | (e0 = op(e1, e3)) | (e0 = op(e1, e2)) | (e0 = op(e1, e1)) | (e0 = op(e1, e0))) & ((e4 = op(e4, e0)) | (e4 = op(e3, e0)) | (e4 = op(e2, e0)) | (e4 = op(e1, e0)) | (op(e0, e0) = e4)) & ((e4 = op(e0, e4)) | (e4 = op(e0, e3)) | (e4 = op(e0, e2)) | (e4 = op(e0, e1)) | (op(e0, e0) = e4)) & ((e3 = op(e4, e0)) | (e3 = op(e3, e0)) | (e3 = op(e2, e0)) | (e3 = op(e1, e0)) | (op(e0, e0) = e3)) & ((e3 = op(e0, e4)) | (e3 = op(e0, e3)) | (e3 = op(e0, e2)) | (e3 = op(e0, e1)) | (op(e0, e0) = e3)) & ((e2 = op(e4, e0)) | (e2 = op(e3, e0)) | (e2 = op(e2, e0)) | (e2 = op(e1, e0)) | (op(e0, e0) = e2)) & ((e2 = op(e0, e4)) | (e2 = op(e0, e3)) | (e2 = op(e0, e2)) | (e2 = op(e0, e1)) | (op(e0, e0) = e2)) & ((e1 = op(e4, e0)) | (e1 = op(e3, e0)) | (e1 = op(e2, e0)) | (e1 = op(e1, e0)) | (op(e0, e0) = e1)) & ((e1 = op(e0, e4)) | (e1 = op(e0, e3)) | (e1 = op(e0, e2)) | (e1 = op(e0, e1)) | (op(e0, e0) = e1)) & ((e0 = op(e4, e0)) | (e0 = op(e3, e0)) | (e0 = op(e2, e0)) | (e0 = op(e1, e0)) | (e0 = op(e0, e0))) & ((e0 = op(e0, e4)) | (e0 = op(e0, e3)) | (e0 = op(e0, e2)) | (e0 = op(e0, e1)) | (e0 = op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG191+1.p', ax2)).
fof(f1184, plain, (spl4_121 | spl4_96 | spl4_71 | spl4_46 | spl4_21), inference(avatar_split_clause, [], [f282, f696, f801, f906, f1011, f1116])).
fof(f282, plain, ((e0 = op(e4, e0)) | (e0 = op(e3, e0)) | (e0 = op(e2, e0)) | (e0 = op(e1, e0)) | (e0 = op(e0, e0))), inference(cnf_transformation, [], [f2])).
fof(f1183, plain, (spl4_122 | spl4_117 | spl4_112 | spl4_107 | spl4_102), inference(avatar_split_clause, [], [f283, f1036, f1057, f1078, f1099, f1120])).
fof(f283, plain, ((e1 = op(e0, e4)) | (e1 = op(e0, e3)) | (e1 = op(e0, e2)) | (e1 = op(e0, e1)) | (op(e0, e0) = e1)), inference(cnf_transformation, [], [f2])).
fof(f1182, plain, (spl4_122 | spl4_97 | spl4_72 | spl4_47 | spl4_22), inference(avatar_split_clause, [], [f284, f700, f805, f910, f1015, f1120])).
fof(f284, plain, ((e1 = op(e4, e0)) | (e1 = op(e3, e0)) | (e1 = op(e2, e0)) | (e1 = op(e1, e0)) | (op(e0, e0) = e1)), inference(cnf_transformation, [], [f2])).
fof(f1181, plain, (spl4_123 | spl4_118 | spl4_113 | spl4_108 | spl4_103), inference(avatar_split_clause, [], [f285, f1040, f1061, f1082, f1103, f1124])).
fof(f285, plain, ((e2 = op(e0, e4)) | (e2 = op(e0, e3)) | (e2 = op(e0, e2)) | (e2 = op(e0, e1)) | (op(e0, e0) = e2)), inference(cnf_transformation, [], [f2])).
fof(f1180, plain, (spl4_123 | spl4_98 | spl4_73 | spl4_48 | spl4_23), inference(avatar_split_clause, [], [f286, f704, f809, f914, f1019, f1124])).
fof(f286, plain, ((e2 = op(e4, e0)) | (e2 = op(e3, e0)) | (e2 = op(e2, e0)) | (e2 = op(e1, e0)) | (op(e0, e0) = e2)), inference(cnf_transformation, [], [f2])).
fof(f1179, plain, (spl4_124 | spl4_119 | spl4_114 | spl4_109 | spl4_104), inference(avatar_split_clause, [], [f287, f1044, f1065, f1086, f1107, f1128])).
fof(f287, plain, ((e3 = op(e0, e4)) | (e3 = op(e0, e3)) | (e3 = op(e0, e2)) | (e3 = op(e0, e1)) | (op(e0, e0) = e3)), inference(cnf_transformation, [], [f2])).
fof(f1178, plain, (spl4_124 | spl4_99 | spl4_74 | spl4_49 | spl4_24), inference(avatar_split_clause, [], [f288, f708, f813, f918, f1023, f1128])).
fof(f288, plain, ((e3 = op(e4, e0)) | (e3 = op(e3, e0)) | (e3 = op(e2, e0)) | (e3 = op(e1, e0)) | (op(e0, e0) = e3)), inference(cnf_transformation, [], [f2])).
fof(f1177, plain, (spl4_125 | spl4_120 | spl4_115 | spl4_110 | spl4_105), inference(avatar_split_clause, [], [f289, f1048, f1069, f1090, f1111, f1132])).
fof(f289, plain, ((e4 = op(e0, e4)) | (e4 = op(e0, e3)) | (e4 = op(e0, e2)) | (e4 = op(e0, e1)) | (op(e0, e0) = e4)), inference(cnf_transformation, [], [f2])).
fof(f1176, plain, (spl4_125 | spl4_100 | spl4_75 | spl4_50 | spl4_25), inference(avatar_split_clause, [], [f290, f712, f817, f922, f1027, f1132])).
fof(f290, plain, ((e4 = op(e4, e0)) | (e4 = op(e3, e0)) | (e4 = op(e2, e0)) | (e4 = op(e1, e0)) | (op(e0, e0) = e4)), inference(cnf_transformation, [], [f2])).
fof(f1175, plain, (spl4_96 | spl4_91 | spl4_86 | spl4_81 | spl4_76), inference(avatar_split_clause, [], [f291, f927, f948, f969, f990, f1011])).
fof(f291, plain, ((e0 = op(e1, e4)) | (e0 = op(e1, e3)) | (e0 = op(e1, e2)) | (e0 = op(e1, e1)) | (e0 = op(e1, e0))), inference(cnf_transformation, [], [f2])).
fof(f1174, plain, (spl4_116 | spl4_91 | spl4_66 | spl4_41 | spl4_16), inference(avatar_split_clause, [], [f292, f675, f780, f885, f990, f1095])).
fof(f292, plain, ((e0 = op(e4, e1)) | (e0 = op(e3, e1)) | (e0 = op(e2, e1)) | (e0 = op(e1, e1)) | (e0 = op(e0, e1))), inference(cnf_transformation, [], [f2])).
fof(f1173, plain, (spl4_97 | spl4_92 | spl4_87 | spl4_82 | spl4_77), inference(avatar_split_clause, [], [f293, f931, f952, f973, f994, f1015])).
fof(f293, plain, ((e1 = op(e1, e4)) | (e1 = op(e1, e3)) | (e1 = op(e1, e2)) | (e1 = op(e1, e1)) | (e1 = op(e1, e0))), inference(cnf_transformation, [], [f2])).
fof(f1172, plain, (spl4_117 | spl4_92 | spl4_67 | spl4_42 | spl4_17), inference(avatar_split_clause, [], [f294, f679, f784, f889, f994, f1099])).
fof(f294, plain, ((e1 = op(e4, e1)) | (e1 = op(e3, e1)) | (e1 = op(e2, e1)) | (e1 = op(e1, e1)) | (e1 = op(e0, e1))), inference(cnf_transformation, [], [f2])).
fof(f1171, plain, (spl4_98 | spl4_93 | spl4_88 | spl4_83 | spl4_78), inference(avatar_split_clause, [], [f295, f935, f956, f977, f998, f1019])).
fof(f295, plain, ((e2 = op(e1, e4)) | (e2 = op(e1, e3)) | (e2 = op(e1, e2)) | (e2 = op(e1, e1)) | (e2 = op(e1, e0))), inference(cnf_transformation, [], [f2])).
fof(f1170, plain, (spl4_118 | spl4_93 | spl4_68 | spl4_43 | spl4_18), inference(avatar_split_clause, [], [f296, f683, f788, f893, f998, f1103])).
fof(f296, plain, ((e2 = op(e4, e1)) | (e2 = op(e3, e1)) | (e2 = op(e2, e1)) | (e2 = op(e1, e1)) | (e2 = op(e0, e1))), inference(cnf_transformation, [], [f2])).
fof(f1169, plain, (spl4_99 | spl4_94 | spl4_89 | spl4_84 | spl4_79), inference(avatar_split_clause, [], [f297, f939, f960, f981, f1002, f1023])).
fof(f297, plain, ((e3 = op(e1, e4)) | (e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, e0))), inference(cnf_transformation, [], [f2])).
fof(f1168, plain, (spl4_119 | spl4_94 | spl4_69 | spl4_44 | spl4_19), inference(avatar_split_clause, [], [f298, f687, f792, f897, f1002, f1107])).
fof(f298, plain, ((e3 = op(e4, e1)) | (e3 = op(e3, e1)) | (e3 = op(e2, e1)) | (e3 = op(e1, e1)) | (e3 = op(e0, e1))), inference(cnf_transformation, [], [f2])).
fof(f1167, plain, (spl4_100 | spl4_95 | spl4_90 | spl4_85 | spl4_80), inference(avatar_split_clause, [], [f299, f943, f964, f985, f1006, f1027])).
fof(f299, plain, ((e4 = op(e1, e4)) | (e4 = op(e1, e3)) | (e4 = op(e1, e2)) | (e4 = op(e1, e1)) | (e4 = op(e1, e0))), inference(cnf_transformation, [], [f2])).
fof(f1166, plain, (spl4_120 | spl4_95 | spl4_70 | spl4_45 | spl4_20), inference(avatar_split_clause, [], [f300, f691, f796, f901, f1006, f1111])).
fof(f300, plain, ((e4 = op(e4, e1)) | (e4 = op(e3, e1)) | (e4 = op(e2, e1)) | (e4 = op(e1, e1)) | (e4 = op(e0, e1))), inference(cnf_transformation, [], [f2])).
fof(f1164, plain, (spl4_111 | spl4_86 | spl4_61 | spl4_36 | spl4_11), inference(avatar_split_clause, [], [f302, f654, f759, f864, f969, f1074])).
fof(f302, plain, ((e0 = op(e4, e2)) | (e0 = op(e3, e2)) | (e0 = op(e2, e2)) | (e0 = op(e1, e2)) | (e0 = op(e0, e2))), inference(cnf_transformation, [], [f2])).
fof(f1163, plain, (spl4_72 | spl4_67 | spl4_62 | spl4_57 | spl4_52), inference(avatar_split_clause, [], [f303, f826, f847, f868, f889, f910])).
fof(f303, plain, ((e1 = op(e2, e4)) | (e1 = op(e2, e3)) | (e1 = op(e2, e2)) | (e1 = op(e2, e1)) | (e1 = op(e2, e0))), inference(cnf_transformation, [], [f2])).
fof(f1162, plain, (spl4_112 | spl4_87 | spl4_62 | spl4_37 | spl4_12), inference(avatar_split_clause, [], [f304, f658, f763, f868, f973, f1078])).
fof(f304, plain, ((e1 = op(e4, e2)) | (e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(e0, e2))), inference(cnf_transformation, [], [f2])).
fof(f1161, plain, (spl4_73 | spl4_68 | spl4_63 | spl4_58 | spl4_53), inference(avatar_split_clause, [], [f305, f830, f851, f872, f893, f914])).
fof(f305, plain, ((e2 = op(e2, e4)) | (e2 = op(e2, e3)) | (e2 = op(e2, e2)) | (e2 = op(e2, e1)) | (e2 = op(e2, e0))), inference(cnf_transformation, [], [f2])).
fof(f1159, plain, (spl4_74 | spl4_69 | spl4_64 | spl4_59 | spl4_54), inference(avatar_split_clause, [], [f307, f834, f855, f876, f897, f918])).
fof(f307, plain, ((e3 = op(e2, e4)) | (e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, e0))), inference(cnf_transformation, [], [f2])).
fof(f1158, plain, (spl4_114 | spl4_89 | spl4_64 | spl4_39 | spl4_14), inference(avatar_split_clause, [], [f308, f666, f771, f876, f981, f1086])).
fof(f308, plain, ((e3 = op(e4, e2)) | (e3 = op(e3, e2)) | (e3 = op(e2, e2)) | (e3 = op(e1, e2)) | (e3 = op(e0, e2))), inference(cnf_transformation, [], [f2])).
fof(f1157, plain, (spl4_75 | spl4_70 | spl4_65 | spl4_60 | spl4_55), inference(avatar_split_clause, [], [f309, f838, f859, f880, f901, f922])).
fof(f309, plain, ((e4 = op(e2, e4)) | (e4 = op(e2, e3)) | (e4 = op(e2, e2)) | (e4 = op(e2, e1)) | (e4 = op(e2, e0))), inference(cnf_transformation, [], [f2])).
fof(f1156, plain, (spl4_115 | spl4_90 | spl4_65 | spl4_40 | spl4_15), inference(avatar_split_clause, [], [f310, f670, f775, f880, f985, f1090])).
fof(f310, plain, ((e4 = op(e4, e2)) | (e4 = op(e3, e2)) | (e4 = op(e2, e2)) | (e4 = op(e1, e2)) | (e4 = op(e0, e2))), inference(cnf_transformation, [], [f2])).
fof(f1155, plain, (spl4_46 | spl4_41 | spl4_36 | spl4_31 | spl4_26), inference(avatar_split_clause, [], [f311, f717, f738, f759, f780, f801])).
fof(f311, plain, ((e0 = op(e3, e4)) | (e0 = op(e3, e3)) | (e0 = op(e3, e2)) | (e0 = op(e3, e1)) | (e0 = op(e3, e0))), inference(cnf_transformation, [], [f2])).
fof(f1154, plain, (spl4_106 | spl4_81 | spl4_56 | spl4_31 | spl4_6), inference(avatar_split_clause, [], [f312, f633, f738, f843, f948, f1053])).
fof(f312, plain, ((e0 = op(e4, e3)) | (e0 = op(e3, e3)) | (e0 = op(e2, e3)) | (e0 = op(e1, e3)) | (e0 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f1153, plain, (spl4_47 | spl4_42 | spl4_37 | spl4_32 | spl4_27), inference(avatar_split_clause, [], [f313, f721, f742, f763, f784, f805])).
fof(f313, plain, ((e1 = op(e3, e4)) | (e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, e0))), inference(cnf_transformation, [], [f2])).
fof(f1152, plain, (spl4_107 | spl4_82 | spl4_57 | spl4_32 | spl4_7), inference(avatar_split_clause, [], [f314, f637, f742, f847, f952, f1057])).
fof(f314, plain, ((e1 = op(e4, e3)) | (e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f1151, plain, (spl4_48 | spl4_43 | spl4_38 | spl4_33 | spl4_28), inference(avatar_split_clause, [], [f315, f725, f746, f767, f788, f809])).
fof(f315, plain, ((e2 = op(e3, e4)) | (e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, e0))), inference(cnf_transformation, [], [f2])).
fof(f1150, plain, (spl4_108 | spl4_83 | spl4_58 | spl4_33 | spl4_8), inference(avatar_split_clause, [], [f316, f641, f746, f851, f956, f1061])).
fof(f316, plain, ((e2 = op(e4, e3)) | (e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f1149, plain, (spl4_49 | spl4_44 | spl4_39 | spl4_34 | spl4_29), inference(avatar_split_clause, [], [f317, f729, f750, f771, f792, f813])).
fof(f317, plain, ((e3 = op(e3, e4)) | (e3 = op(e3, e3)) | (e3 = op(e3, e2)) | (e3 = op(e3, e1)) | (e3 = op(e3, e0))), inference(cnf_transformation, [], [f2])).
fof(f1148, plain, (spl4_109 | spl4_84 | spl4_59 | spl4_34 | spl4_9), inference(avatar_split_clause, [], [f318, f645, f750, f855, f960, f1065])).
fof(f318, plain, ((e3 = op(e4, e3)) | (e3 = op(e3, e3)) | (e3 = op(e2, e3)) | (e3 = op(e1, e3)) | (e3 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f1147, plain, (spl4_50 | spl4_45 | spl4_40 | spl4_35 | spl4_30), inference(avatar_split_clause, [], [f319, f733, f754, f775, f796, f817])).
fof(f319, plain, ((e4 = op(e3, e4)) | (e4 = op(e3, e3)) | (e4 = op(e3, e2)) | (e4 = op(e3, e1)) | (e4 = op(e3, e0))), inference(cnf_transformation, [], [f2])).
fof(f1146, plain, (spl4_110 | spl4_85 | spl4_60 | spl4_35 | spl4_10), inference(avatar_split_clause, [], [f320, f649, f754, f859, f964, f1069])).
fof(f320, plain, ((e4 = op(e4, e3)) | (e4 = op(e3, e3)) | (e4 = op(e2, e3)) | (e4 = op(e1, e3)) | (e4 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f1145, plain, (spl4_21 | spl4_16 | spl4_11 | spl4_6 | spl4_1), inference(avatar_split_clause, [], [f321, f612, f633, f654, f675, f696])).
fof(f321, plain, ((e0 = op(e4, e4)) | (e0 = op(e4, e3)) | (e0 = op(e4, e2)) | (e0 = op(e4, e1)) | (e0 = op(e4, e0))), inference(cnf_transformation, [], [f2])).
fof(f1143, plain, (spl4_22 | spl4_17 | spl4_12 | spl4_7 | spl4_2), inference(avatar_split_clause, [], [f323, f616, f637, f658, f679, f700])).
fof(f323, plain, ((e1 = op(e4, e4)) | (e1 = op(e4, e3)) | (e1 = op(e4, e2)) | (e1 = op(e4, e1)) | (e1 = op(e4, e0))), inference(cnf_transformation, [], [f2])).
fof(f1142, plain, (spl4_102 | spl4_77 | spl4_52 | spl4_27 | spl4_2), inference(avatar_split_clause, [], [f324, f616, f721, f826, f931, f1036])).
fof(f324, plain, ((e1 = op(e4, e4)) | (e1 = op(e3, e4)) | (e1 = op(e2, e4)) | (e1 = op(e1, e4)) | (e1 = op(e0, e4))), inference(cnf_transformation, [], [f2])).
fof(f1141, plain, (spl4_23 | spl4_18 | spl4_13 | spl4_8 | spl4_3), inference(avatar_split_clause, [], [f325, f620, f641, f662, f683, f704])).
fof(f325, plain, ((e2 = op(e4, e4)) | (e2 = op(e4, e3)) | (e2 = op(e4, e2)) | (e2 = op(e4, e1)) | (e2 = op(e4, e0))), inference(cnf_transformation, [], [f2])).
fof(f1140, plain, (spl4_103 | spl4_78 | spl4_53 | spl4_28 | spl4_3), inference(avatar_split_clause, [], [f326, f620, f725, f830, f935, f1040])).
fof(f326, plain, ((e2 = op(e4, e4)) | (e2 = op(e3, e4)) | (e2 = op(e2, e4)) | (e2 = op(e1, e4)) | (e2 = op(e0, e4))), inference(cnf_transformation, [], [f2])).
fof(f1139, plain, (spl4_24 | spl4_19 | spl4_14 | spl4_9 | spl4_4), inference(avatar_split_clause, [], [f327, f624, f645, f666, f687, f708])).
fof(f327, plain, ((e3 = op(e4, e4)) | (e3 = op(e4, e3)) | (e3 = op(e4, e2)) | (e3 = op(e4, e1)) | (e3 = op(e4, e0))), inference(cnf_transformation, [], [f2])).
fof(f1138, plain, (spl4_104 | spl4_79 | spl4_54 | spl4_29 | spl4_4), inference(avatar_split_clause, [], [f328, f624, f729, f834, f939, f1044])).
fof(f328, plain, ((e3 = op(e4, e4)) | (e3 = op(e3, e4)) | (e3 = op(e2, e4)) | (e3 = op(e1, e4)) | (e3 = op(e0, e4))), inference(cnf_transformation, [], [f2])).
fof(f1137, plain, (spl4_25 | spl4_20 | spl4_15 | spl4_10 | spl4_5), inference(avatar_split_clause, [], [f329, f628, f649, f670, f691, f712])).
fof(f329, plain, ((e4 = op(e4, e4)) | (e4 = op(e4, e3)) | (e4 = op(e4, e2)) | (e4 = op(e4, e1)) | (e4 = op(e4, e0))), inference(cnf_transformation, [], [f2])).
fof(f1136, plain, (spl4_105 | spl4_80 | spl4_55 | spl4_30 | spl4_5), inference(avatar_split_clause, [], [f330, f628, f733, f838, f943, f1048])).
fof(f330, plain, ((e4 = op(e4, e4)) | (e4 = op(e3, e4)) | (e4 = op(e2, e4)) | (e4 = op(e1, e4)) | (e4 = op(e0, e4))), inference(cnf_transformation, [], [f2])).
fof(f1135, plain, (spl4_121 | spl4_122 | spl4_123 | spl4_124 | spl4_125), inference(avatar_split_clause, [], [f256, f1132, f1128, f1124, f1120, f1116])).
fof(f256, plain, ((op(e0, e0) = e4) | (op(e0, e0) = e3) | (op(e0, e0) = e2) | (op(e0, e0) = e1) | (e0 = op(e0, e0))), inference(cnf_transformation, [], [f1])).
fof(f1, plain, (((e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | (e1 = op(e4, e4)) | (e0 = op(e4, e4))) & ((e4 = op(e4, e3)) | (e3 = op(e4, e3)) | (e2 = op(e4, e3)) | (e1 = op(e4, e3)) | (e0 = op(e4, e3))) & ((e4 = op(e4, e2)) | (e3 = op(e4, e2)) | (e2 = op(e4, e2)) | (e1 = op(e4, e2)) | (e0 = op(e4, e2))) & ((e4 = op(e4, e1)) | (e3 = op(e4, e1)) | (e2 = op(e4, e1)) | (e1 = op(e4, e1)) | (e0 = op(e4, e1))) & ((e4 = op(e4, e0)) | (e3 = op(e4, e0)) | (e2 = op(e4, e0)) | (e1 = op(e4, e0)) | (e0 = op(e4, e0))) & ((e4 = op(e3, e4)) | (e3 = op(e3, e4)) | (e2 = op(e3, e4)) | (e1 = op(e3, e4)) | (e0 = op(e3, e4))) & ((e4 = op(e3, e3)) | (e3 = op(e3, e3)) | (e2 = op(e3, e3)) | (e1 = op(e3, e3)) | (e0 = op(e3, e3))) & ((e4 = op(e3, e2)) | (e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (e0 = op(e3, e2))) & ((e4 = op(e3, e1)) | (e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (e0 = op(e3, e1))) & ((e4 = op(e3, e0)) | (e3 = op(e3, e0)) | (e2 = op(e3, e0)) | (e1 = op(e3, e0)) | (e0 = op(e3, e0))) & ((e4 = op(e2, e4)) | (e3 = op(e2, e4)) | (e2 = op(e2, e4)) | (e1 = op(e2, e4)) | (e0 = op(e2, e4))) & ((e4 = op(e2, e3)) | (e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))) & ((e4 = op(e2, e2)) | (e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (e0 = op(e2, e2))) & ((e4 = op(e2, e1)) | (e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))) & ((e4 = op(e2, e0)) | (e3 = op(e2, e0)) | (e2 = op(e2, e0)) | (e1 = op(e2, e0)) | (e0 = op(e2, e0))) & ((e4 = op(e1, e4)) | (e3 = op(e1, e4)) | (e2 = op(e1, e4)) | (e1 = op(e1, e4)) | (e0 = op(e1, e4))) & ((e4 = op(e1, e3)) | (e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (e0 = op(e1, e3))) & ((e4 = op(e1, e2)) | (e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (e0 = op(e1, e2))) & ((e4 = op(e1, e1)) | (e3 = op(e1, e1)) | (e2 = op(e1, e1)) | (e1 = op(e1, e1)) | (e0 = op(e1, e1))) & ((e4 = op(e1, e0)) | (e3 = op(e1, e0)) | (e2 = op(e1, e0)) | (e1 = op(e1, e0)) | (e0 = op(e1, e0))) & ((e4 = op(e0, e4)) | (e3 = op(e0, e4)) | (e2 = op(e0, e4)) | (e1 = op(e0, e4)) | (e0 = op(e0, e4))) & ((e4 = op(e0, e3)) | (e3 = op(e0, e3)) | (e2 = op(e0, e3)) | (e1 = op(e0, e3)) | (e0 = op(e0, e3))) & ((e4 = op(e0, e2)) | (e3 = op(e0, e2)) | (e2 = op(e0, e2)) | (e1 = op(e0, e2)) | (e0 = op(e0, e2))) & ((e4 = op(e0, e1)) | (e3 = op(e0, e1)) | (e2 = op(e0, e1)) | (e1 = op(e0, e1)) | (e0 = op(e0, e1))) & ((op(e0, e0) = e4) | (op(e0, e0) = e3) | (op(e0, e0) = e2) | (op(e0, e0) = e1) | (e0 = op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG191+1.p', ax1)).
fof(f1114, plain, (spl4_116 | spl4_117 | spl4_118 | spl4_119 | spl4_120), inference(avatar_split_clause, [], [f257, f1111, f1107, f1103, f1099, f1095])).
fof(f257, plain, ((e4 = op(e0, e1)) | (e3 = op(e0, e1)) | (e2 = op(e0, e1)) | (e1 = op(e0, e1)) | (e0 = op(e0, e1))), inference(cnf_transformation, [], [f1])).
fof(f1093, plain, (spl4_111 | spl4_112 | spl4_113 | spl4_114 | spl4_115), inference(avatar_split_clause, [], [f258, f1090, f1086, f1082, f1078, f1074])).
fof(f258, plain, ((e4 = op(e0, e2)) | (e3 = op(e0, e2)) | (e2 = op(e0, e2)) | (e1 = op(e0, e2)) | (e0 = op(e0, e2))), inference(cnf_transformation, [], [f1])).
fof(f1072, plain, (spl4_106 | spl4_107 | spl4_108 | spl4_109 | spl4_110), inference(avatar_split_clause, [], [f259, f1069, f1065, f1061, f1057, f1053])).
fof(f259, plain, ((e4 = op(e0, e3)) | (e3 = op(e0, e3)) | (e2 = op(e0, e3)) | (e1 = op(e0, e3)) | (e0 = op(e0, e3))), inference(cnf_transformation, [], [f1])).
fof(f1051, plain, (spl4_101 | spl4_102 | spl4_103 | spl4_104 | spl4_105), inference(avatar_split_clause, [], [f260, f1048, f1044, f1040, f1036, f1032])).
fof(f260, plain, ((e4 = op(e0, e4)) | (e3 = op(e0, e4)) | (e2 = op(e0, e4)) | (e1 = op(e0, e4)) | (e0 = op(e0, e4))), inference(cnf_transformation, [], [f1])).
fof(f1030, plain, (spl4_96 | spl4_97 | spl4_98 | spl4_99 | spl4_100), inference(avatar_split_clause, [], [f261, f1027, f1023, f1019, f1015, f1011])).
fof(f261, plain, ((e4 = op(e1, e0)) | (e3 = op(e1, e0)) | (e2 = op(e1, e0)) | (e1 = op(e1, e0)) | (e0 = op(e1, e0))), inference(cnf_transformation, [], [f1])).
fof(f1009, plain, (spl4_91 | spl4_92 | spl4_93 | spl4_94 | spl4_95), inference(avatar_split_clause, [], [f262, f1006, f1002, f998, f994, f990])).
fof(f262, plain, ((e4 = op(e1, e1)) | (e3 = op(e1, e1)) | (e2 = op(e1, e1)) | (e1 = op(e1, e1)) | (e0 = op(e1, e1))), inference(cnf_transformation, [], [f1])).
fof(f988, plain, (spl4_86 | spl4_87 | spl4_88 | spl4_89 | spl4_90), inference(avatar_split_clause, [], [f263, f985, f981, f977, f973, f969])).
fof(f263, plain, ((e4 = op(e1, e2)) | (e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (e0 = op(e1, e2))), inference(cnf_transformation, [], [f1])).
fof(f967, plain, (spl4_81 | spl4_82 | spl4_83 | spl4_84 | spl4_85), inference(avatar_split_clause, [], [f264, f964, f960, f956, f952, f948])).
fof(f264, plain, ((e4 = op(e1, e3)) | (e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (e0 = op(e1, e3))), inference(cnf_transformation, [], [f1])).
fof(f946, plain, (spl4_76 | spl4_77 | spl4_78 | spl4_79 | spl4_80), inference(avatar_split_clause, [], [f265, f943, f939, f935, f931, f927])).
fof(f265, plain, ((e4 = op(e1, e4)) | (e3 = op(e1, e4)) | (e2 = op(e1, e4)) | (e1 = op(e1, e4)) | (e0 = op(e1, e4))), inference(cnf_transformation, [], [f1])).
fof(f925, plain, (spl4_71 | spl4_72 | spl4_73 | spl4_74 | spl4_75), inference(avatar_split_clause, [], [f266, f922, f918, f914, f910, f906])).
fof(f266, plain, ((e4 = op(e2, e0)) | (e3 = op(e2, e0)) | (e2 = op(e2, e0)) | (e1 = op(e2, e0)) | (e0 = op(e2, e0))), inference(cnf_transformation, [], [f1])).
fof(f904, plain, (spl4_66 | spl4_67 | spl4_68 | spl4_69 | spl4_70), inference(avatar_split_clause, [], [f267, f901, f897, f893, f889, f885])).
fof(f267, plain, ((e4 = op(e2, e1)) | (e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))), inference(cnf_transformation, [], [f1])).
fof(f883, plain, (spl4_61 | spl4_62 | spl4_63 | spl4_64 | spl4_65), inference(avatar_split_clause, [], [f268, f880, f876, f872, f868, f864])).
fof(f268, plain, ((e4 = op(e2, e2)) | (e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (e0 = op(e2, e2))), inference(cnf_transformation, [], [f1])).
fof(f862, plain, (spl4_56 | spl4_57 | spl4_58 | spl4_59 | spl4_60), inference(avatar_split_clause, [], [f269, f859, f855, f851, f847, f843])).
fof(f269, plain, ((e4 = op(e2, e3)) | (e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))), inference(cnf_transformation, [], [f1])).
fof(f841, plain, (spl4_51 | spl4_52 | spl4_53 | spl4_54 | spl4_55), inference(avatar_split_clause, [], [f270, f838, f834, f830, f826, f822])).
fof(f270, plain, ((e4 = op(e2, e4)) | (e3 = op(e2, e4)) | (e2 = op(e2, e4)) | (e1 = op(e2, e4)) | (e0 = op(e2, e4))), inference(cnf_transformation, [], [f1])).
fof(f820, plain, (spl4_46 | spl4_47 | spl4_48 | spl4_49 | spl4_50), inference(avatar_split_clause, [], [f271, f817, f813, f809, f805, f801])).
fof(f271, plain, ((e4 = op(e3, e0)) | (e3 = op(e3, e0)) | (e2 = op(e3, e0)) | (e1 = op(e3, e0)) | (e0 = op(e3, e0))), inference(cnf_transformation, [], [f1])).
fof(f799, plain, (spl4_41 | spl4_42 | spl4_43 | spl4_44 | spl4_45), inference(avatar_split_clause, [], [f272, f796, f792, f788, f784, f780])).
fof(f272, plain, ((e4 = op(e3, e1)) | (e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (e0 = op(e3, e1))), inference(cnf_transformation, [], [f1])).
fof(f778, plain, (spl4_36 | spl4_37 | spl4_38 | spl4_39 | spl4_40), inference(avatar_split_clause, [], [f273, f775, f771, f767, f763, f759])).
fof(f273, plain, ((e4 = op(e3, e2)) | (e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (e0 = op(e3, e2))), inference(cnf_transformation, [], [f1])).
fof(f757, plain, (spl4_31 | spl4_32 | spl4_33 | spl4_34 | spl4_35), inference(avatar_split_clause, [], [f274, f754, f750, f746, f742, f738])).
fof(f274, plain, ((e4 = op(e3, e3)) | (e3 = op(e3, e3)) | (e2 = op(e3, e3)) | (e1 = op(e3, e3)) | (e0 = op(e3, e3))), inference(cnf_transformation, [], [f1])).
fof(f736, plain, (spl4_26 | spl4_27 | spl4_28 | spl4_29 | spl4_30), inference(avatar_split_clause, [], [f275, f733, f729, f725, f721, f717])).
fof(f275, plain, ((e4 = op(e3, e4)) | (e3 = op(e3, e4)) | (e2 = op(e3, e4)) | (e1 = op(e3, e4)) | (e0 = op(e3, e4))), inference(cnf_transformation, [], [f1])).
fof(f715, plain, (spl4_21 | spl4_22 | spl4_23 | spl4_24 | spl4_25), inference(avatar_split_clause, [], [f276, f712, f708, f704, f700, f696])).
fof(f276, plain, ((e4 = op(e4, e0)) | (e3 = op(e4, e0)) | (e2 = op(e4, e0)) | (e1 = op(e4, e0)) | (e0 = op(e4, e0))), inference(cnf_transformation, [], [f1])).
fof(f694, plain, (spl4_16 | spl4_17 | spl4_18 | spl4_19 | spl4_20), inference(avatar_split_clause, [], [f277, f691, f687, f683, f679, f675])).
fof(f277, plain, ((e4 = op(e4, e1)) | (e3 = op(e4, e1)) | (e2 = op(e4, e1)) | (e1 = op(e4, e1)) | (e0 = op(e4, e1))), inference(cnf_transformation, [], [f1])).
fof(f673, plain, (spl4_11 | spl4_12 | spl4_13 | spl4_14 | spl4_15), inference(avatar_split_clause, [], [f278, f670, f666, f662, f658, f654])).
fof(f278, plain, ((e4 = op(e4, e2)) | (e3 = op(e4, e2)) | (e2 = op(e4, e2)) | (e1 = op(e4, e2)) | (e0 = op(e4, e2))), inference(cnf_transformation, [], [f1])).
fof(f652, plain, (spl4_6 | spl4_7 | spl4_8 | spl4_9 | spl4_10), inference(avatar_split_clause, [], [f279, f649, f645, f641, f637, f633])).
fof(f279, plain, ((e4 = op(e4, e3)) | (e3 = op(e4, e3)) | (e2 = op(e4, e3)) | (e1 = op(e4, e3)) | (e0 = op(e4, e3))), inference(cnf_transformation, [], [f1])).
fof(f631, plain, (spl4_1 | spl4_2 | spl4_3 | spl4_4 | spl4_5), inference(avatar_split_clause, [], [f280, f628, f624, f620, f616, f612])).
fof(f280, plain, ((e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | (e1 = op(e4, e4)) | (e0 = op(e4, e4))), inference(cnf_transformation, [], [f1])).