fof(f4679, plain, $false, inference(avatar_sat_refutation, [], [f605, f622, f639, f656, f673, f690, f707, f724, f741, f758, f792, f809, f826, f843, f860, f861, f862, f863, f864, f865, f866, f867, f868, f869, f870, f871, f872, f873, f874, f875, f876, f877, f878, f879, f880, f881, f882, f883, f885, f886, f887, f888, f889, f890, f891, f897, f898, f900, f905, f906, f908, f913, f916, f921, f922, f923, f924, f930, f931, f932, f937, f938, f939, f940, f945, f946, f947, f948, f953, f954, f955, f956, f961, f962, f963, f964, f970, f971, f972, f977, f978, f979, f980, f985, f986, f987, f988, f993, f994, f995, f996, f1001, f1002, f1003, f1004, f1010, f1011, f1012, f1018, f1019, f1020, f1025, f1026, f1027, f1028, f1033, f1034, f1035, f1036, f1041, f1042, f1043, f1044, f1049, f1050, f1052, f1060, f1068, f1074, f1076, f1081, f1082, f1083, f1084, f1089, f1090, f1091, f1092, f1098, f1099, f1100, f1105, f1106, f1107, f1108, f1113, f1114, f1115, f1116, f1121, f1123, f1124, f1129, f1130, f1131, f1132, f1138, f1139, f1140, f1146, f1147, f1153, f1154, f1155, f1156, f1161, f1162, f1163, f1164, f1169, f1170, f1171, f1172, f1177, f1178, f1179, f1180, f1186, f1188, f1193, f1194, f1195, f1196, f1201, f1202, f1203, f1204, f1209, f1210, f1212, f1217, f1218, f1220, f1228, f1233, f1234, f1236, f1241, f1242, f1243, f1244, f1249, f1250, f1251, f1252, f1257, f1258, f1259, f1260, f1266, f1267, f1268, f1274, f1275, f1276, f1281, f1282, f1283, f1284, f1289, f1290, f1291, f1292, f1297, f1298, f1299, f1300, f1305, f1306, f1308, f1314, f1316, f1321, f1322, f1324, f1329, f1330, f1332, f1337, f1338, f1339, f1340, f1345, f1347, f1348, f1354, f1355, f1356, f1361, f1362, f1363, f1364, f1369, f1370, f1372, f1377, f1378, f1380, f1385, f1386, f1388, f1396, f1405, f1410, f1415, f1420, f1429, f1434, f1439, f1444, f1453, f1458, f1463, f1468, f1472, f1477, f1482, f1487, f1492, f1493, f1494, f1495, f1496, f1505, f1519, f1538, f1548, f1553, f1572, f1583, f1588, f1598, f1603, f1609, f1614, f1615, f1616, f1626, f1630, f1655, f1664, f1667, f1686, f1687, f1699, f1701, f1708, f1711, f1716, f1718, f1724, f1741, f1744, f1755, f1758, f1770, f1780, f1783, f1793, f1802, f1805, f1807, f1816, f1837, f1850, f1853, f1867, f1882, f1903, f1920, f1923, f1935, f1959, f1961, f1985, f2010, f2029, f2034, f2067, f2070, f2073, f2078, f2088, f2093, f2099, f2125, f2151, f2168, f2175, f2176, f2187, f2189, f2218, f2231, f2243, f2247, f2289, f2303, f2325, f2338, f2339, f2374, f2441, f2456, f2466, f2499, f2500, f2528, f2534, f2536, f2556, f2561, f2574, f2583, f2591, f2592, f2594, f2610, f2618, f2622, f2627, f2634, f2638, f2666, f2677, f2685, f2719, f2733, f2734, f2735, f2752, f2763, f2766, f2775, f2789, f2801, f2804, f2864, f2873, f2878, f2899, f2914, f2921, f2923, f2928, f2932, f2951, f2967, f2975, f2998, f3019, f3034, f3067, f3086, f3100, f3110, f3116, f3118, f3119, f3121, f3134, f3137, f3139, f3152, f3154, f3177, f3181, f3192, f3194, f3227, f3234, f3247, f3248, f3261, f3266, f3268, f3294, f3304, f3305, f3315, f3324, f3340, f3372, f3375, f3396, f3397, f3407, f3424, f3440, f3459, f3473, f3474, f3478, f3485, f3486, f3487, f3496, f3497, f3514, f3528, f3540, f3551, f3557, f3558, f3568, f3582, f3583, f3601, f3610, f3612, f3621, f3622, f3627, f3642, f3664, f3688, f3698, f3732, f3733, f3738, f3756, f3760, f3771, f3772, f3774, f3782, f3783, f3789, f3806, f3816, f3817, f3831, f3832, f3842, f3854, f3857, f3860, f3870, f3871, f3878, f3879, f3887, f3897, f3898, f3899, f3908, f3920, f3933, f3938, f3945, f3954, f3955, f3958, f3978, f3989, f3991, f3993, f4008, f4009, f4011, f4016, f4022, f4029, f4044, f4056, f4058, f4078, f4084, f4089, f4096, f4098, f4100, f4102, f4106, f4108, f4113, f4115, f4128, f4130, f4144, f4145, f4148, f4156, f4162, f4163, f4169, f4177, f4191, f4206, f4213, f4247, f4249, f4250, f4257, f4258, f4270, f4291, f4295, f4302, f4305, f4317, f4333, f4342, f4360, f4361, f4376, f4386, f4396, f4397, f4424, f4435, f4443, f4455, f4461, f4465, f4471, f4480, f4481, f4492, f4500, f4501, f4506, f4507, f4522, f4537, f4541, f4552, f4558, f4588, f4589, f4591, f4592, f4599, f4600, f4601, f4609, f4615, f4623, f4634, f4653, f4668])).
fof(f4668, plain, (~ spl66_19 | ~ spl66_23), inference(avatar_split_clause, [], [f4660, f683, f666])).
fof(f666, plain, (spl66_19 <=> (e2 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_19])])).
fof(f683, plain, (spl66_23 <=> (e2 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_23])])).
fof(f4660, plain, (~ (e2 = op(e2, e3)) | ~ spl66_23), inference(backward_demodulation, [], [f276, f685])).
fof(f685, plain, ((e2 = op(e2, e2)) | ~ spl66_23), inference(avatar_component_clause, [], [f683])).
fof(f276, plain, ~ (op(e2, e2) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f3, plain, (~ (op(e3, e2) = op(e3, e3)) & ~ (op(e3, e1) = op(e3, e3)) & ~ (op(e3, e1) = op(e3, e2)) & ~ (op(e3, e0) = op(e3, e3)) & ~ (op(e3, e0) = op(e3, e2)) & ~ (op(e3, e0) = op(e3, e1)) & ~ (op(e2, e2) = op(e2, e3)) & ~ (op(e2, e1) = op(e2, e3)) & ~ (op(e2, e1) = op(e2, e2)) & ~ (op(e2, e0) = op(e2, e3)) & ~ (op(e2, e0) = op(e2, e2)) & ~ (op(e2, e0) = op(e2, e1)) & ~ (op(e1, e2) = op(e1, e3)) & ~ (op(e1, e1) = op(e1, e3)) & ~ (op(e1, e1) = op(e1, e2)) & ~ (op(e1, e0) = op(e1, e3)) & ~ (op(e1, e0) = op(e1, e2)) & ~ (op(e1, e0) = op(e1, e1)) & ~ (op(e0, e2) = op(e0, e3)) & ~ (op(e0, e1) = op(e0, e3)) & ~ (op(e0, e1) = op(e0, e2)) & ~ (op(e0, e0) = op(e0, e3)) & ~ (op(e0, e0) = op(e0, e2)) & ~ (op(e0, e0) = op(e0, e1)) & ~ (op(e2, e3) = op(e3, e3)) & ~ (op(e1, e3) = op(e3, e3)) & ~ (op(e1, e3) = op(e2, e3)) & ~ (op(e0, e3) = op(e3, e3)) & ~ (op(e0, e3) = op(e2, e3)) & ~ (op(e0, e3) = op(e1, e3)) & ~ (op(e2, e2) = op(e3, e2)) & ~ (op(e1, e2) = op(e3, e2)) & ~ (op(e1, e2) = op(e2, e2)) & ~ (op(e0, e2) = op(e3, e2)) & ~ (op(e0, e2) = op(e2, e2)) & ~ (op(e0, e2) = op(e1, e2)) & ~ (op(e2, e1) = op(e3, e1)) & ~ (op(e1, e1) = op(e3, e1)) & ~ (op(e1, e1) = op(e2, e1)) & ~ (op(e0, e1) = op(e3, e1)) & ~ (op(e0, e1) = op(e2, e1)) & ~ (op(e0, e1) = op(e1, e1)) & ~ (op(e2, e0) = op(e3, e0)) & ~ (op(e1, e0) = op(e3, e0)) & ~ (op(e1, e0) = op(e2, e0)) & ~ (op(e0, e0) = op(e3, e0)) & ~ (op(e0, e0) = op(e2, e0)) & ~ (op(e0, e0) = op(e1, e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG140+1.p', ax3)).
fof(f4653, plain, (~ spl66_28 | ~ spl66_32), inference(avatar_split_clause, [], [f4649, f721, f704])).
fof(f704, plain, (spl66_28 <=> (e3 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_28])])).
fof(f721, plain, (spl66_32 <=> (e3 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_32])])).
fof(f4649, plain, (~ (e3 = op(e2, e1)) | ~ spl66_32), inference(backward_demodulation, [], [f271, f723])).
fof(f723, plain, ((e3 = op(e2, e0)) | ~ spl66_32), inference(avatar_component_clause, [], [f721])).
fof(f271, plain, ~ (op(e2, e0) = op(e2, e1)), inference(cnf_transformation, [], [f3])).
fof(f4634, plain, (~ spl66_33 | ~ spl66_41), inference(avatar_split_clause, [], [f4626, f760, f726])).
fof(f726, plain, (spl66_33 <=> (e0 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_33])])).
fof(f760, plain, (spl66_41 <=> (e0 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_41])])).
fof(f4626, plain, (~ (e0 = op(e1, e3)) | ~ spl66_41), inference(backward_demodulation, [], [f269, f762])).
fof(f762, plain, ((e0 = op(e1, e1)) | ~ spl66_41), inference(avatar_component_clause, [], [f760])).
fof(f269, plain, ~ (op(e1, e1) = op(e1, e3)), inference(cnf_transformation, [], [f3])).
fof(f4623, plain, (~ spl66_39 | ~ spl66_47), inference(avatar_split_clause, [], [f4619, f785, f751])).
fof(f751, plain, (spl66_39 <=> (e2 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_39])])).
fof(f785, plain, (spl66_47 <=> (e2 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_47])])).
fof(f4619, plain, (~ (e2 = op(e1, e2)) | ~ spl66_47), inference(backward_demodulation, [], [f266, f787])).
fof(f787, plain, ((e2 = op(e1, e0)) | ~ spl66_47), inference(avatar_component_clause, [], [f785])).
fof(f266, plain, ~ (op(e1, e0) = op(e1, e2)), inference(cnf_transformation, [], [f3])).
fof(f4615, plain, (~ spl66_19 | ~ spl66_51), inference(avatar_split_clause, [], [f4612, f802, f666])).
fof(f802, plain, (spl66_51 <=> (e2 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_51])])).
fof(f4612, plain, (~ (e2 = op(e2, e3)) | ~ spl66_51), inference(backward_demodulation, [], [f254, f804])).
fof(f804, plain, ((e2 = op(e0, e3)) | ~ spl66_51), inference(avatar_component_clause, [], [f802])).
fof(f254, plain, ~ (op(e0, e3) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f4609, plain, (~ spl66_21 | ~ spl66_53), inference(avatar_split_clause, [], [f4603, f811, f675])).
fof(f675, plain, (spl66_21 <=> (e0 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_21])])).
fof(f811, plain, (spl66_53 <=> (e0 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_53])])).
fof(f4603, plain, (~ (e0 = op(e2, e2)) | ~ spl66_53), inference(backward_demodulation, [], [f248, f813])).
fof(f813, plain, ((e0 = op(e0, e2)) | ~ spl66_53), inference(avatar_component_clause, [], [f811])).
fof(f248, plain, ~ (op(e0, e2) = op(e2, e2)), inference(cnf_transformation, [], [f3])).
fof(f4601, plain, (~ spl66_50 | ~ spl66_60 | spl66_141), inference(avatar_split_clause, [], [f4598, f1460, f840, f798])).
fof(f798, plain, (spl66_50 <=> (e1 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_50])])).
fof(f840, plain, (spl66_60 <=> (e3 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_60])])).
fof(f1460, plain, (spl66_141 <=> (e1 = op(e0, op(e0, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl66_141])])).
fof(f4598, plain, (~ (e1 = op(e0, e3)) | (~ spl66_60 | spl66_141)), inference(backward_demodulation, [], [f1462, f842])).
fof(f842, plain, ((e3 = op(e0, e1)) | ~ spl66_60), inference(avatar_component_clause, [], [f840])).
fof(f1462, plain, (~ (e1 = op(e0, op(e0, e1))) | spl66_141), inference(avatar_component_clause, [], [f1460])).
fof(f4600, plain, (~ spl66_56 | ~ spl66_60), inference(avatar_split_clause, [], [f4596, f840, f823])).
fof(f823, plain, (spl66_56 <=> (e3 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_56])])).
fof(f4596, plain, (~ (e3 = op(e0, e2)) | ~ spl66_60), inference(backward_demodulation, [], [f262, f842])).
fof(f262, plain, ~ (op(e0, e1) = op(e0, e2)), inference(cnf_transformation, [], [f3])).
fof(f4599, plain, (~ spl66_28 | ~ spl66_60), inference(avatar_split_clause, [], [f4595, f840, f704])).
fof(f4595, plain, (~ (e3 = op(e2, e1)) | ~ spl66_60), inference(backward_demodulation, [], [f242, f842])).
fof(f242, plain, ~ (op(e0, e1) = op(e2, e1)), inference(cnf_transformation, [], [f3])).
fof(f4592, plain, (~ spl66_40 | ~ spl66_62 | ~ spl66_162 | spl66_170), inference(avatar_split_clause, [], [f4586, f1611, f1569, f849, f755])).
fof(f755, plain, (spl66_40 <=> (e3 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_40])])).
fof(f849, plain, (spl66_62 <=> (op(e0, e0) = e1)), introduced(avatar_definition, [new_symbols(naming, [spl66_62])])).
fof(f1569, plain, (spl66_162 <=> (e2 = op(op(e0, e0), e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_162])])).
fof(f1611, plain, (spl66_170 <=> (e3 = op(op(e0, e0), op(op(e0, e0), e0)))), introduced(avatar_definition, [new_symbols(naming, [spl66_170])])).
fof(f4586, plain, (~ (e3 = op(e1, e2)) | (~ spl66_62 | ~ spl66_162 | spl66_170)), inference(backward_demodulation, [], [f4573, f851])).
fof(f851, plain, ((op(e0, e0) = e1) | ~ spl66_62), inference(avatar_component_clause, [], [f849])).
fof(f4573, plain, (~ (e3 = op(op(e0, e0), e2)) | (~ spl66_162 | spl66_170)), inference(backward_demodulation, [], [f1613, f1570])).
fof(f1570, plain, ((e2 = op(op(e0, e0), e0)) | ~ spl66_162), inference(avatar_component_clause, [], [f1569])).
fof(f1613, plain, (~ (e3 = op(op(e0, e0), op(op(e0, e0), e0))) | spl66_170), inference(avatar_component_clause, [], [f1611])).
fof(f4591, plain, (spl66_47 | ~ spl66_62 | ~ spl66_162), inference(avatar_split_clause, [], [f4583, f1569, f849, f785])).
fof(f4583, plain, ((e2 = op(e1, e0)) | (~ spl66_62 | ~ spl66_162)), inference(backward_demodulation, [], [f1570, f851])).
fof(f4589, plain, (~ spl66_50 | ~ spl66_62), inference(avatar_split_clause, [], [f4579, f849, f798])).
fof(f4579, plain, (~ (e1 = op(e0, e3)) | ~ spl66_62), inference(backward_demodulation, [], [f261, f851])).
fof(f261, plain, ~ (op(e0, e0) = op(e0, e3)), inference(cnf_transformation, [], [f3])).
fof(f4588, plain, (~ spl66_30 | ~ spl66_62), inference(avatar_split_clause, [], [f4576, f849, f713])).
fof(f713, plain, (spl66_30 <=> (e1 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_30])])).
fof(f4576, plain, (~ (e1 = op(e2, e0)) | ~ spl66_62), inference(backward_demodulation, [], [f236, f851])).
fof(f236, plain, ~ (op(e0, e0) = op(e2, e0)), inference(cnf_transformation, [], [f3])).
fof(f4558, plain, (~ spl66_28 | ~ spl66_30 | ~ spl66_63 | spl66_170), inference(avatar_contradiction_clause, [], [f4557])).
fof(f4557, plain, ($false | (~ spl66_28 | ~ spl66_30 | ~ spl66_63 | spl66_170)), inference(subsumption_resolution, [], [f4556, f706])).
fof(f706, plain, ((e3 = op(e2, e1)) | ~ spl66_28), inference(avatar_component_clause, [], [f704])).
fof(f4556, plain, (~ (e3 = op(e2, e1)) | (~ spl66_30 | ~ spl66_63 | spl66_170)), inference(forward_demodulation, [], [f4555, f715])).
fof(f715, plain, ((e1 = op(e2, e0)) | ~ spl66_30), inference(avatar_component_clause, [], [f713])).
fof(f4555, plain, (~ (e3 = op(e2, op(e2, e0))) | (~ spl66_63 | spl66_170)), inference(forward_demodulation, [], [f1613, f855])).
fof(f855, plain, ((op(e0, e0) = e2) | ~ spl66_63), inference(avatar_component_clause, [], [f853])).
fof(f853, plain, (spl66_63 <=> (op(e0, e0) = e2)), introduced(avatar_definition, [new_symbols(naming, [spl66_63])])).
fof(f4552, plain, (~ spl66_21 | ~ spl66_50 | ~ spl66_56 | spl66_157), inference(avatar_contradiction_clause, [], [f4551])).
fof(f4551, plain, ($false | (~ spl66_21 | ~ spl66_50 | ~ spl66_56 | spl66_157)), inference(subsumption_resolution, [], [f4550, f800])).
fof(f800, plain, ((e1 = op(e0, e3)) | ~ spl66_50), inference(avatar_component_clause, [], [f798])).
fof(f4550, plain, (~ (e1 = op(e0, e3)) | (~ spl66_21 | ~ spl66_56 | spl66_157)), inference(forward_demodulation, [], [f4549, f825])).
fof(f825, plain, ((e3 = op(e0, e2)) | ~ spl66_56), inference(avatar_component_clause, [], [f823])).
fof(f4549, plain, (~ (e1 = op(e0, op(e0, e2))) | (~ spl66_21 | spl66_157)), inference(forward_demodulation, [], [f1547, f677])).
fof(f677, plain, ((e0 = op(e2, e2)) | ~ spl66_21), inference(avatar_component_clause, [], [f675])).
fof(f1547, plain, (~ (e1 = op(op(e2, e2), op(op(e2, e2), e2))) | spl66_157), inference(avatar_component_clause, [], [f1545])).
fof(f1545, plain, (spl66_157 <=> (e1 = op(op(e2, e2), op(op(e2, e2), e2)))), introduced(avatar_definition, [new_symbols(naming, [spl66_157])])).
fof(f4541, plain, (~ spl66_7 | ~ spl66_11), inference(avatar_split_clause, [], [f4538, f632, f615])).
fof(f615, plain, (spl66_7 <=> (e2 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_7])])).
fof(f632, plain, (spl66_11 <=> (e2 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_11])])).
fof(f4538, plain, (~ (e2 = op(e3, e2)) | ~ spl66_11), inference(backward_demodulation, [], [f280, f634])).
fof(f634, plain, ((e2 = op(e3, e1)) | ~ spl66_11), inference(avatar_component_clause, [], [f632])).
fof(f280, plain, ~ (op(e3, e1) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f4537, plain, (~ spl66_9 | ~ spl66_13), inference(avatar_split_clause, [], [f4531, f641, f624])).
fof(f624, plain, (spl66_9 <=> (e0 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_9])])).
fof(f641, plain, (spl66_13 <=> (e0 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_13])])).
fof(f4531, plain, (~ (e0 = op(e3, e1)) | ~ spl66_13), inference(backward_demodulation, [], [f277, f643])).
fof(f643, plain, ((e0 = op(e3, e0)) | ~ spl66_13), inference(avatar_component_clause, [], [f641])).
fof(f277, plain, ~ (op(e3, e0) = op(e3, e1)), inference(cnf_transformation, [], [f3])).
fof(f4522, plain, (~ spl66_18 | ~ spl66_30), inference(avatar_split_clause, [], [f4515, f713, f662])).
fof(f662, plain, (spl66_18 <=> (e1 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_18])])).
fof(f4515, plain, (~ (e1 = op(e2, e3)) | ~ spl66_30), inference(backward_demodulation, [], [f273, f715])).
fof(f273, plain, ~ (op(e2, e0) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f4507, plain, (~ spl66_35 | ~ spl66_39), inference(avatar_split_clause, [], [f4503, f751, f734])).
fof(f734, plain, (spl66_35 <=> (e2 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_35])])).
fof(f4503, plain, (~ (e2 = op(e1, e3)) | ~ spl66_39), inference(backward_demodulation, [], [f270, f753])).
fof(f753, plain, ((e2 = op(e1, e2)) | ~ spl66_39), inference(avatar_component_clause, [], [f751])).
fof(f270, plain, ~ (op(e1, e2) = op(e1, e3)), inference(cnf_transformation, [], [f3])).
fof(f4506, plain, (~ spl66_7 | ~ spl66_39), inference(avatar_split_clause, [], [f4502, f751, f615])).
fof(f4502, plain, (~ (e2 = op(e3, e2)) | ~ spl66_39), inference(backward_demodulation, [], [f251, f753])).
fof(f251, plain, ~ (op(e1, e2) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f4501, plain, (spl66_33 | ~ spl66_48 | ~ spl66_137), inference(avatar_split_clause, [], [f4498, f1441, f789, f726])).
fof(f789, plain, (spl66_48 <=> (e3 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_48])])).
fof(f1441, plain, (spl66_137 <=> (e0 = op(e1, op(e1, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl66_137])])).
fof(f4498, plain, ((e0 = op(e1, e3)) | (~ spl66_48 | ~ spl66_137)), inference(backward_demodulation, [], [f1442, f791])).
fof(f791, plain, ((e3 = op(e1, e0)) | ~ spl66_48), inference(avatar_component_clause, [], [f789])).
fof(f1442, plain, ((e0 = op(e1, op(e1, e0))) | ~ spl66_137), inference(avatar_component_clause, [], [f1441])).
fof(f4500, plain, (~ spl66_40 | ~ spl66_48), inference(avatar_split_clause, [], [f4496, f789, f755])).
fof(f4496, plain, (~ (e3 = op(e1, e2)) | ~ spl66_48), inference(backward_demodulation, [], [f266, f791])).
fof(f4492, plain, (~ spl66_18 | ~ spl66_50), inference(avatar_split_clause, [], [f4488, f798, f662])).
fof(f4488, plain, (~ (e1 = op(e2, e3)) | ~ spl66_50), inference(backward_demodulation, [], [f254, f800])).
fof(f4481, plain, (~ spl66_49 | ~ spl66_57), inference(avatar_split_clause, [], [f4475, f828, f794])).
fof(f794, plain, (spl66_49 <=> (e0 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_49])])).
fof(f828, plain, (spl66_57 <=> (e0 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_57])])).
fof(f4475, plain, (~ (e0 = op(e0, e3)) | ~ spl66_57), inference(backward_demodulation, [], [f263, f830])).
fof(f830, plain, ((e0 = op(e0, e1)) | ~ spl66_57), inference(avatar_component_clause, [], [f828])).
fof(f263, plain, ~ (op(e0, e1) = op(e0, e3)), inference(cnf_transformation, [], [f3])).
fof(f4480, plain, (~ spl66_9 | ~ spl66_57), inference(avatar_split_clause, [], [f4473, f828, f624])).
fof(f4473, plain, (~ (e0 = op(e3, e1)) | ~ spl66_57), inference(backward_demodulation, [], [f243, f830])).
fof(f243, plain, ~ (op(e0, e1) = op(e3, e1)), inference(cnf_transformation, [], [f3])).
fof(f4471, plain, (spl66_56 | ~ spl66_21 | ~ spl66_150), inference(avatar_split_clause, [], [f4470, f1511, f675, f823])).
fof(f1511, plain, (spl66_150 <=> (e3 = op(op(e2, e2), e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_150])])).
fof(f4470, plain, ((e3 = op(e0, e2)) | (~ spl66_21 | ~ spl66_150)), inference(forward_demodulation, [], [f1512, f677])).
fof(f1512, plain, ((e3 = op(op(e2, e2), e2)) | ~ spl66_150), inference(avatar_component_clause, [], [f1511])).
fof(f4465, plain, (~ spl66_60 | ~ spl66_21 | ~ spl66_154 | spl66_168), inference(avatar_split_clause, [], [f4434, f1600, f1530, f675, f840])).
fof(f1530, plain, (spl66_154 <=> (e1 = op(op(e2, e2), e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_154])])).
fof(f1600, plain, (spl66_168 <=> (e3 = op(op(e2, e2), op(op(e2, e2), e2)))), introduced(avatar_definition, [new_symbols(naming, [spl66_168])])).
fof(f4434, plain, (~ (e3 = op(e0, e1)) | (~ spl66_21 | ~ spl66_154 | spl66_168)), inference(backward_demodulation, [], [f4346, f677])).
fof(f4346, plain, (~ (e3 = op(op(e2, e2), e1)) | (~ spl66_154 | spl66_168)), inference(backward_demodulation, [], [f1602, f1531])).
fof(f1531, plain, ((e1 = op(op(e2, e2), e2)) | ~ spl66_154), inference(avatar_component_clause, [], [f1530])).
fof(f1602, plain, (~ (e3 = op(op(e2, e2), op(op(e2, e2), e2))) | spl66_168), inference(avatar_component_clause, [], [f1600])).
fof(f4461, plain, (~ spl66_51 | ~ spl66_63), inference(avatar_split_clause, [], [f4460, f853, f802])).
fof(f4460, plain, (~ (e2 = op(e0, e3)) | ~ spl66_63), inference(forward_demodulation, [], [f261, f855])).
fof(f4455, plain, (~ spl66_30 | ~ spl66_63 | spl66_167), inference(avatar_split_clause, [], [f4359, f1595, f853, f713])).
fof(f1595, plain, (spl66_167 <=> (e1 = op(op(e0, e0), e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_167])])).
fof(f4359, plain, (~ (e1 = op(e2, e0)) | (~ spl66_63 | spl66_167)), inference(backward_demodulation, [], [f1597, f855])).
fof(f1597, plain, (~ (e1 = op(op(e0, e0), e0)) | spl66_167), inference(avatar_component_clause, [], [f1595])).
fof(f4443, plain, (~ spl66_10 | ~ spl66_14), inference(avatar_split_clause, [], [f4438, f645, f628])).
fof(f628, plain, (spl66_10 <=> (e1 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_10])])).
fof(f645, plain, (spl66_14 <=> (e1 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_14])])).
fof(f4438, plain, (~ (e1 = op(e3, e1)) | ~ spl66_14), inference(backward_demodulation, [], [f277, f647])).
fof(f647, plain, ((e1 = op(e3, e0)) | ~ spl66_14), inference(avatar_component_clause, [], [f645])).
fof(f4435, plain, (~ spl66_5 | ~ spl66_21), inference(avatar_split_clause, [], [f4425, f675, f607])).
fof(f607, plain, (spl66_5 <=> (e0 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_5])])).
fof(f4425, plain, (~ (e0 = op(e3, e2)) | ~ spl66_21), inference(backward_demodulation, [], [f252, f677])).
fof(f252, plain, ~ (op(e2, e2) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f4424, plain, (~ spl66_23 | ~ spl66_27), inference(avatar_split_clause, [], [f4421, f700, f683])).
fof(f700, plain, (spl66_27 <=> (e2 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_27])])).
fof(f4421, plain, (~ (e2 = op(e2, e2)) | ~ spl66_27), inference(backward_demodulation, [], [f274, f702])).
fof(f702, plain, ((e2 = op(e2, e1)) | ~ spl66_27), inference(avatar_component_clause, [], [f700])).
fof(f274, plain, ~ (op(e2, e1) = op(e2, e2)), inference(cnf_transformation, [], [f3])).
fof(f4397, plain, (~ spl66_38 | ~ spl66_42), inference(avatar_split_clause, [], [f4389, f764, f747])).
fof(f747, plain, (spl66_38 <=> (e1 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_38])])).
fof(f764, plain, (spl66_42 <=> (e1 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_42])])).
fof(f4389, plain, (~ (e1 = op(e1, e2)) | ~ spl66_42), inference(backward_demodulation, [], [f268, f766])).
fof(f766, plain, ((e1 = op(e1, e1)) | ~ spl66_42), inference(avatar_component_clause, [], [f764])).
fof(f268, plain, ~ (op(e1, e1) = op(e1, e2)), inference(cnf_transformation, [], [f3])).
fof(f4396, plain, (~ spl66_10 | ~ spl66_42), inference(avatar_split_clause, [], [f4388, f764, f628])).
fof(f4388, plain, (~ (e1 = op(e3, e1)) | ~ spl66_42), inference(backward_demodulation, [], [f245, f766])).
fof(f245, plain, ~ (op(e1, e1) = op(e3, e1)), inference(cnf_transformation, [], [f3])).
fof(f4386, plain, (~ spl66_41 | ~ spl66_45), inference(avatar_split_clause, [], [f4379, f777, f760])).
fof(f777, plain, (spl66_45 <=> (e0 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_45])])).
fof(f4379, plain, (~ (e0 = op(e1, e1)) | ~ spl66_45), inference(backward_demodulation, [], [f265, f779])).
fof(f779, plain, ((e0 = op(e1, e0)) | ~ spl66_45), inference(avatar_component_clause, [], [f777])).
fof(f265, plain, ~ (op(e1, e0) = op(e1, e1)), inference(cnf_transformation, [], [f3])).
fof(f4376, plain, (~ spl66_38 | ~ spl66_54), inference(avatar_split_clause, [], [f4370, f815, f747])).
fof(f815, plain, (spl66_54 <=> (e1 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_54])])).
fof(f4370, plain, (~ (e1 = op(e1, e2)) | ~ spl66_54), inference(backward_demodulation, [], [f247, f817])).
fof(f817, plain, ((e1 = op(e0, e2)) | ~ spl66_54), inference(avatar_component_clause, [], [f815])).
fof(f247, plain, ~ (op(e0, e2) = op(e1, e2)), inference(cnf_transformation, [], [f3])).
fof(f4361, plain, (~ spl66_59 | ~ spl66_63), inference(avatar_split_clause, [], [f4352, f853, f836])).
fof(f836, plain, (spl66_59 <=> (e2 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_59])])).
fof(f4352, plain, (~ (e2 = op(e0, e1)) | ~ spl66_63), inference(backward_demodulation, [], [f259, f855])).
fof(f259, plain, ~ (op(e0, e0) = op(e0, e1)), inference(cnf_transformation, [], [f3])).
fof(f4360, plain, (~ spl66_15 | ~ spl66_63), inference(avatar_split_clause, [], [f4351, f853, f649])).
fof(f649, plain, (spl66_15 <=> (e2 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_15])])).
fof(f4351, plain, (~ (e2 = op(e3, e0)) | ~ spl66_63), inference(backward_demodulation, [], [f237, f855])).
fof(f237, plain, ~ (op(e0, e0) = op(e3, e0)), inference(cnf_transformation, [], [f3])).
fof(f4342, plain, (~ spl66_64 | ~ spl66_49 | spl66_139), inference(avatar_split_clause, [], [f4269, f1450, f794, f857])).
fof(f857, plain, (spl66_64 <=> (op(e0, e0) = e3)), introduced(avatar_definition, [new_symbols(naming, [spl66_64])])).
fof(f1450, plain, (spl66_139 <=> (e3 = op(e0, op(e0, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl66_139])])).
fof(f4269, plain, (~ (op(e0, e0) = e3) | (~ spl66_49 | spl66_139)), inference(backward_demodulation, [], [f1452, f796])).
fof(f796, plain, ((e0 = op(e0, e3)) | ~ spl66_49), inference(avatar_component_clause, [], [f794])).
fof(f1452, plain, (~ (e3 = op(e0, op(e0, e3))) | spl66_139), inference(avatar_component_clause, [], [f1450])).
fof(f4333, plain, (~ spl66_12 | ~ spl66_4), inference(avatar_split_clause, [], [f4103, f602, f636])).
fof(f636, plain, (spl66_12 <=> (e3 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_12])])).
fof(f602, plain, (spl66_4 <=> (e3 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_4])])).
fof(f4103, plain, (~ (e3 = op(e3, e1)) | ~ spl66_4), inference(forward_demodulation, [], [f281, f604])).
fof(f604, plain, ((e3 = op(e3, e3)) | ~ spl66_4), inference(avatar_component_clause, [], [f602])).
fof(f281, plain, ~ (op(e3, e1) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f4317, plain, (spl66_18 | ~ spl66_28 | ~ spl66_131), inference(avatar_split_clause, [], [f4316, f1412, f704, f662])).
fof(f1412, plain, (spl66_131 <=> (e1 = op(e2, op(e2, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl66_131])])).
fof(f4316, plain, ((e1 = op(e2, e3)) | (~ spl66_28 | ~ spl66_131)), inference(backward_demodulation, [], [f1413, f706])).
fof(f1413, plain, ((e1 = op(e2, op(e2, e1))) | ~ spl66_131), inference(avatar_component_clause, [], [f1412])).
fof(f4305, plain, (~ spl66_19 | ~ spl66_35), inference(avatar_split_clause, [], [f4303, f734, f666])).
fof(f4303, plain, (~ (e2 = op(e2, e3)) | ~ spl66_35), inference(backward_demodulation, [], [f256, f736])).
fof(f736, plain, ((e2 = op(e1, e3)) | ~ spl66_35), inference(avatar_component_clause, [], [f734])).
fof(f256, plain, ~ (op(e1, e3) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f4302, plain, (~ spl66_22 | ~ spl66_38), inference(avatar_split_clause, [], [f4298, f747, f679])).
fof(f679, plain, (spl66_22 <=> (e1 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_22])])).
fof(f4298, plain, (~ (e1 = op(e2, e2)) | ~ spl66_38), inference(backward_demodulation, [], [f250, f749])).
fof(f749, plain, ((e1 = op(e1, e2)) | ~ spl66_38), inference(avatar_component_clause, [], [f747])).
fof(f250, plain, ~ (op(e1, e2) = op(e2, e2)), inference(cnf_transformation, [], [f3])).
fof(f4295, plain, (~ spl66_41 | ~ spl66_59 | spl66_155), inference(avatar_contradiction_clause, [], [f4294])).
fof(f4294, plain, ($false | (~ spl66_41 | ~ spl66_59 | spl66_155)), inference(subsumption_resolution, [], [f4288, f838])).
fof(f838, plain, ((e2 = op(e0, e1)) | ~ spl66_59), inference(avatar_component_clause, [], [f836])).
fof(f4288, plain, (~ (e2 = op(e0, e1)) | (~ spl66_41 | spl66_155)), inference(backward_demodulation, [], [f1537, f762])).
fof(f1537, plain, (~ (e2 = op(op(e1, e1), e1)) | spl66_155), inference(avatar_component_clause, [], [f1535])).
fof(f1535, plain, (spl66_155 <=> (e2 = op(op(e1, e1), e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_155])])).
fof(f4291, plain, (~ spl66_25 | ~ spl66_41), inference(avatar_split_clause, [], [f4281, f760, f692])).
fof(f692, plain, (spl66_25 <=> (e0 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_25])])).
fof(f4281, plain, (~ (e0 = op(e2, e1)) | ~ spl66_41), inference(backward_demodulation, [], [f244, f762])).
fof(f244, plain, ~ (op(e1, e1) = op(e2, e1)), inference(cnf_transformation, [], [f3])).
fof(f4270, plain, (~ spl66_33 | ~ spl66_49), inference(avatar_split_clause, [], [f4264, f794, f726])).
fof(f4264, plain, (~ (e0 = op(e1, e3)) | ~ spl66_49), inference(backward_demodulation, [], [f253, f796])).
fof(f253, plain, ~ (op(e0, e3) = op(e1, e3)), inference(cnf_transformation, [], [f3])).
fof(f4258, plain, (~ spl66_55 | ~ spl66_59), inference(avatar_split_clause, [], [f4253, f836, f819])).
fof(f819, plain, (spl66_55 <=> (e2 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_55])])).
fof(f4253, plain, (~ (e2 = op(e0, e2)) | ~ spl66_59), inference(backward_demodulation, [], [f262, f838])).
fof(f4257, plain, (~ spl66_43 | ~ spl66_59), inference(avatar_split_clause, [], [f4251, f836, f768])).
fof(f768, plain, (spl66_43 <=> (e2 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_43])])).
fof(f4251, plain, (~ (e2 = op(e1, e1)) | ~ spl66_59), inference(backward_demodulation, [], [f241, f838])).
fof(f241, plain, ~ (op(e0, e1) = op(e1, e1)), inference(cnf_transformation, [], [f3])).
fof(f4250, plain, (~ spl66_46 | ~ spl66_62 | spl66_167), inference(avatar_split_clause, [], [f4246, f1595, f849, f781])).
fof(f781, plain, (spl66_46 <=> (e1 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_46])])).
fof(f4246, plain, (~ (e1 = op(e1, e0)) | (~ spl66_62 | spl66_167)), inference(backward_demodulation, [], [f1597, f851])).
fof(f4249, plain, (~ spl66_48 | ~ spl66_62 | spl66_160), inference(avatar_split_clause, [], [f4243, f1559, f849, f789])).
fof(f1559, plain, (spl66_160 <=> (e3 = op(op(e0, e0), e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_160])])).
fof(f4243, plain, (~ (e3 = op(e1, e0)) | (~ spl66_62 | spl66_160)), inference(backward_demodulation, [], [f1561, f851])).
fof(f1561, plain, (~ (e3 = op(op(e0, e0), e0)) | spl66_160), inference(avatar_component_clause, [], [f1559])).
fof(f4247, plain, (~ spl66_46 | ~ spl66_62), inference(avatar_split_clause, [], [f4234, f849, f781])).
fof(f4234, plain, (~ (e1 = op(e1, e0)) | ~ spl66_62), inference(backward_demodulation, [], [f235, f851])).
fof(f235, plain, ~ (op(e0, e0) = op(e1, e0)), inference(cnf_transformation, [], [f3])).
fof(f4213, plain, (~ spl66_7 | ~ spl66_15), inference(avatar_split_clause, [], [f4210, f649, f615])).
fof(f4210, plain, (~ (e2 = op(e3, e2)) | ~ spl66_15), inference(backward_demodulation, [], [f278, f651])).
fof(f651, plain, ((e2 = op(e3, e0)) | ~ spl66_15), inference(avatar_component_clause, [], [f649])).
fof(f278, plain, ~ (op(e3, e0) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f4206, plain, (~ spl66_17 | ~ spl66_25), inference(avatar_split_clause, [], [f4199, f692, f658])).
fof(f658, plain, (spl66_17 <=> (e0 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_17])])).
fof(f4199, plain, (~ (e0 = op(e2, e3)) | ~ spl66_25), inference(backward_demodulation, [], [f275, f694])).
fof(f694, plain, ((e0 = op(e2, e1)) | ~ spl66_25), inference(avatar_component_clause, [], [f692])).
fof(f275, plain, ~ (op(e2, e1) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f4191, plain, (~ spl66_17 | ~ spl66_33), inference(avatar_split_clause, [], [f4186, f726, f658])).
fof(f4186, plain, (~ (e0 = op(e2, e3)) | ~ spl66_33), inference(backward_demodulation, [], [f256, f728])).
fof(f728, plain, ((e0 = op(e1, e3)) | ~ spl66_33), inference(avatar_component_clause, [], [f726])).
fof(f4177, plain, (~ spl66_34 | ~ spl66_46), inference(avatar_split_clause, [], [f4173, f781, f730])).
fof(f730, plain, (spl66_34 <=> (e1 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_34])])).
fof(f4173, plain, (~ (e1 = op(e1, e3)) | ~ spl66_46), inference(backward_demodulation, [], [f267, f783])).
fof(f783, plain, ((e1 = op(e1, e0)) | ~ spl66_46), inference(avatar_component_clause, [], [f781])).
fof(f267, plain, ~ (op(e1, e0) = op(e1, e3)), inference(cnf_transformation, [], [f3])).
fof(f4169, plain, (~ spl66_34 | ~ spl66_50), inference(avatar_split_clause, [], [f4164, f798, f730])).
fof(f4164, plain, (~ (e1 = op(e1, e3)) | ~ spl66_50), inference(backward_demodulation, [], [f253, f800])).
fof(f4163, plain, (~ spl66_51 | ~ spl66_55), inference(avatar_split_clause, [], [f4159, f819, f802])).
fof(f4159, plain, (~ (e2 = op(e0, e3)) | ~ spl66_55), inference(backward_demodulation, [], [f264, f821])).
fof(f821, plain, ((e2 = op(e0, e2)) | ~ spl66_55), inference(avatar_component_clause, [], [f819])).
fof(f264, plain, ~ (op(e0, e2) = op(e0, e3)), inference(cnf_transformation, [], [f3])).
fof(f4162, plain, (~ spl66_7 | ~ spl66_55), inference(avatar_split_clause, [], [f4158, f819, f615])).
fof(f4158, plain, (~ (e2 = op(e3, e2)) | ~ spl66_55), inference(backward_demodulation, [], [f249, f821])).
fof(f249, plain, ~ (op(e0, e2) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f4156, plain, (spl66_50 | ~ spl66_60 | ~ spl66_141), inference(avatar_split_clause, [], [f4153, f1460, f840, f798])).
fof(f4153, plain, ((e1 = op(e0, e3)) | (~ spl66_60 | ~ spl66_141)), inference(backward_demodulation, [], [f1461, f842])).
fof(f1461, plain, ((e1 = op(e0, op(e0, e1))) | ~ spl66_141), inference(avatar_component_clause, [], [f1460])).
fof(f4148, plain, (~ spl66_61 | ~ spl66_167), inference(avatar_contradiction_clause, [], [f4147])).
fof(f4147, plain, ($false | (~ spl66_61 | ~ spl66_167)), inference(subsumption_resolution, [], [f4146, f283])).
fof(f283, plain, ~ (e0 = e1), inference(cnf_transformation, [], [f4])).
fof(f4, plain, (~ (e2 = e3) & ~ (e1 = e3) & ~ (e1 = e2) & ~ (e0 = e3) & ~ (e0 = e2) & ~ (e0 = e1)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG140+1.p', ax4)).
fof(f4146, plain, ((e0 = e1) | (~ spl66_61 | ~ spl66_167)), inference(forward_demodulation, [], [f4141, f847])).
fof(f847, plain, ((e0 = op(e0, e0)) | ~ spl66_61), inference(avatar_component_clause, [], [f845])).
fof(f845, plain, (spl66_61 <=> (e0 = op(e0, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_61])])).
fof(f4141, plain, ((op(e0, e0) = e1) | (~ spl66_61 | ~ spl66_167)), inference(backward_demodulation, [], [f1596, f847])).
fof(f1596, plain, ((e1 = op(op(e0, e0), e0)) | ~ spl66_167), inference(avatar_component_clause, [], [f1595])).
fof(f4145, plain, (~ spl66_57 | ~ spl66_61), inference(avatar_split_clause, [], [f4134, f845, f828])).
fof(f4134, plain, (~ (e0 = op(e0, e1)) | ~ spl66_61), inference(backward_demodulation, [], [f259, f847])).
fof(f4144, plain, (~ spl66_13 | ~ spl66_61), inference(avatar_split_clause, [], [f4133, f845, f641])).
fof(f4133, plain, (~ (e0 = op(e3, e0)) | ~ spl66_61), inference(backward_demodulation, [], [f237, f847])).
fof(f4130, plain, (spl66_40 | ~ spl66_22 | ~ spl66_150), inference(avatar_split_clause, [], [f4129, f1511, f679, f755])).
fof(f4129, plain, ((e3 = op(e1, e2)) | (~ spl66_22 | ~ spl66_150)), inference(forward_demodulation, [], [f1512, f681])).
fof(f681, plain, ((e1 = op(e2, e2)) | ~ spl66_22), inference(avatar_component_clause, [], [f679])).
fof(f4128, plain, (~ spl66_37 | ~ spl66_22 | spl66_161), inference(avatar_split_clause, [], [f4127, f1564, f679, f743])).
fof(f743, plain, (spl66_37 <=> (e0 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_37])])).
fof(f1564, plain, (spl66_161 <=> (e0 = op(op(e2, e2), e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_161])])).
fof(f4127, plain, (~ (e0 = op(e1, e2)) | (~ spl66_22 | spl66_161)), inference(forward_demodulation, [], [f1566, f681])).
fof(f1566, plain, (~ (e0 = op(op(e2, e2), e2)) | spl66_161), inference(avatar_component_clause, [], [f1564])).
fof(f4115, plain, (~ spl66_39 | ~ spl66_43), inference(avatar_split_clause, [], [f4114, f768, f751])).
fof(f4114, plain, (~ (e2 = op(e1, e2)) | ~ spl66_43), inference(forward_demodulation, [], [f268, f770])).
fof(f770, plain, ((e2 = op(e1, e1)) | ~ spl66_43), inference(avatar_component_clause, [], [f768])).
fof(f4113, plain, (~ spl66_38 | ~ spl66_43 | spl66_136), inference(avatar_split_clause, [], [f4026, f1436, f768, f747])).
fof(f1436, plain, (spl66_136 <=> (e1 = op(e1, op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl66_136])])).
fof(f4026, plain, (~ (e1 = op(e1, e2)) | (~ spl66_43 | spl66_136)), inference(backward_demodulation, [], [f1438, f770])).
fof(f1438, plain, (~ (e1 = op(e1, op(e1, e1))) | spl66_136), inference(avatar_component_clause, [], [f1436])).
fof(f4108, plain, (~ spl66_25 | ~ spl66_43 | spl66_166), inference(avatar_split_clause, [], [f4028, f1590, f768, f692])).
fof(f1590, plain, (spl66_166 <=> (e0 = op(op(e1, e1), e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_166])])).
fof(f4028, plain, (~ (e0 = op(e2, e1)) | (~ spl66_43 | spl66_166)), inference(backward_demodulation, [], [f1592, f770])).
fof(f1592, plain, (~ (e0 = op(op(e1, e1), e1)) | spl66_166), inference(avatar_component_clause, [], [f1590])).
fof(f4106, plain, (~ spl66_8 | ~ spl66_4), inference(avatar_split_clause, [], [f4105, f602, f619])).
fof(f619, plain, (spl66_8 <=> (e3 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_8])])).
fof(f4105, plain, (~ (e3 = op(e3, e2)) | ~ spl66_4), inference(forward_demodulation, [], [f282, f604])).
fof(f282, plain, ~ (op(e3, e2) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f4102, plain, (~ spl66_16 | ~ spl66_4), inference(avatar_split_clause, [], [f4101, f602, f653])).
fof(f653, plain, (spl66_16 <=> (e3 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_16])])).
fof(f4101, plain, (~ (e3 = op(e3, e0)) | ~ spl66_4), inference(forward_demodulation, [], [f279, f604])).
fof(f279, plain, ~ (op(e3, e0) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f4100, plain, (~ spl66_20 | ~ spl66_4), inference(avatar_split_clause, [], [f4099, f602, f670])).
fof(f670, plain, (spl66_20 <=> (e3 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_20])])).
fof(f4099, plain, (~ (e3 = op(e2, e3)) | ~ spl66_4), inference(forward_demodulation, [], [f258, f604])).
fof(f258, plain, ~ (op(e2, e3) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f4098, plain, (~ spl66_36 | ~ spl66_4), inference(avatar_split_clause, [], [f4097, f602, f738])).
fof(f738, plain, (spl66_36 <=> (e3 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_36])])).
fof(f4097, plain, (~ (e3 = op(e1, e3)) | ~ spl66_4), inference(forward_demodulation, [], [f257, f604])).
fof(f257, plain, ~ (op(e1, e3) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f4096, plain, (~ spl66_52 | ~ spl66_4), inference(avatar_split_clause, [], [f4095, f602, f806])).
fof(f806, plain, (spl66_52 <=> (e3 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_52])])).
fof(f4095, plain, (~ (e3 = op(e0, e3)) | ~ spl66_4), inference(forward_demodulation, [], [f255, f604])).
fof(f255, plain, ~ (op(e0, e3) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f4089, plain, (spl66_1 | ~ spl66_4 | ~ spl66_158), inference(avatar_split_clause, [], [f4088, f1550, f602, f590])).
fof(f590, plain, (spl66_1 <=> (e0 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_1])])).
fof(f1550, plain, (spl66_158 <=> (e0 = op(op(e3, e3), e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_158])])).
fof(f4088, plain, ((e0 = op(e3, e3)) | (~ spl66_4 | ~ spl66_158)), inference(forward_demodulation, [], [f1551, f604])).
fof(f1551, plain, ((e0 = op(op(e3, e3), e3)) | ~ spl66_158), inference(avatar_component_clause, [], [f1550])).
fof(f4084, plain, (spl66_3 | ~ spl66_4 | ~ spl66_148), inference(avatar_split_clause, [], [f4083, f1502, f602, f598])).
fof(f598, plain, (spl66_3 <=> (e2 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_3])])).
fof(f1502, plain, (spl66_148 <=> (e2 = op(op(e3, e3), e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_148])])).
fof(f4083, plain, ((e2 = op(e3, e3)) | (~ spl66_4 | ~ spl66_148)), inference(forward_demodulation, [], [f1503, f604])).
fof(f1503, plain, ((e2 = op(op(e3, e3), e3)) | ~ spl66_148), inference(avatar_component_clause, [], [f1502])).
fof(f4078, plain, (~ spl66_1 | ~ spl66_4), inference(avatar_contradiction_clause, [], [f4077])).
fof(f4077, plain, ($false | (~ spl66_1 | ~ spl66_4)), inference(subsumption_resolution, [], [f4076, f285])).
fof(f285, plain, ~ (e0 = e3), inference(cnf_transformation, [], [f4])).
fof(f4076, plain, ((e0 = e3) | (~ spl66_1 | ~ spl66_4)), inference(forward_demodulation, [], [f604, f592])).
fof(f592, plain, ((e0 = op(e3, e3)) | ~ spl66_1), inference(avatar_component_clause, [], [f590])).
fof(f4058, plain, (~ spl66_22 | ~ spl66_48 | ~ spl66_161 | spl66_168), inference(avatar_contradiction_clause, [], [f4057])).
fof(f4057, plain, ($false | (~ spl66_22 | ~ spl66_48 | ~ spl66_161 | spl66_168)), inference(subsumption_resolution, [], [f4055, f791])).
fof(f4055, plain, (~ (e3 = op(e1, e0)) | (~ spl66_22 | ~ spl66_161 | spl66_168)), inference(backward_demodulation, [], [f3975, f681])).
fof(f3975, plain, (~ (e3 = op(op(e2, e2), e0)) | (~ spl66_161 | spl66_168)), inference(backward_demodulation, [], [f1602, f1565])).
fof(f1565, plain, ((e0 = op(op(e2, e2), e2)) | ~ spl66_161), inference(avatar_component_clause, [], [f1564])).
fof(f4056, plain, (~ spl66_18 | ~ spl66_22), inference(avatar_split_clause, [], [f4048, f679, f662])).
fof(f4048, plain, (~ (e1 = op(e2, e3)) | ~ spl66_22), inference(backward_demodulation, [], [f276, f681])).
fof(f4044, plain, (~ spl66_23 | ~ spl66_31), inference(avatar_split_clause, [], [f4040, f717, f683])).
fof(f717, plain, (spl66_31 <=> (e2 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_31])])).
fof(f4040, plain, (~ (e2 = op(e2, e2)) | ~ spl66_31), inference(backward_demodulation, [], [f272, f719])).
fof(f719, plain, ((e2 = op(e2, e0)) | ~ spl66_31), inference(avatar_component_clause, [], [f717])).
fof(f272, plain, ~ (op(e2, e0) = op(e2, e2)), inference(cnf_transformation, [], [f3])).
fof(f4029, plain, (~ spl66_11 | ~ spl66_43), inference(avatar_split_clause, [], [f4023, f768, f632])).
fof(f4023, plain, (~ (e2 = op(e3, e1)) | ~ spl66_43), inference(backward_demodulation, [], [f245, f770])).
fof(f4022, plain, (~ spl66_36 | ~ spl66_48), inference(avatar_split_clause, [], [f4020, f789, f738])).
fof(f4020, plain, (~ (e3 = op(e1, e3)) | ~ spl66_48), inference(backward_demodulation, [], [f267, f791])).
fof(f4016, plain, (~ spl66_8 | ~ spl66_56), inference(avatar_split_clause, [], [f4014, f823, f619])).
fof(f4014, plain, (~ (e3 = op(e3, e2)) | ~ spl66_56), inference(backward_demodulation, [], [f249, f825])).
fof(f4011, plain, (spl66_46 | ~ spl66_62 | ~ spl66_167), inference(avatar_contradiction_clause, [], [f4010])).
fof(f4010, plain, ($false | (spl66_46 | ~ spl66_62 | ~ spl66_167)), inference(subsumption_resolution, [], [f4005, f782])).
fof(f782, plain, (~ (e1 = op(e1, e0)) | spl66_46), inference(avatar_component_clause, [], [f781])).
fof(f4005, plain, ((e1 = op(e1, e0)) | (~ spl66_62 | ~ spl66_167)), inference(backward_demodulation, [], [f1596, f851])).
fof(f4009, plain, (~ spl66_54 | ~ spl66_62), inference(avatar_split_clause, [], [f4001, f849, f815])).
fof(f4001, plain, (~ (e1 = op(e0, e2)) | ~ spl66_62), inference(backward_demodulation, [], [f260, f851])).
fof(f260, plain, ~ (op(e0, e0) = op(e0, e2)), inference(cnf_transformation, [], [f3])).
fof(f4008, plain, (~ spl66_14 | ~ spl66_62), inference(avatar_split_clause, [], [f4000, f849, f645])).
fof(f4000, plain, (~ (e1 = op(e3, e0)) | ~ spl66_62), inference(backward_demodulation, [], [f237, f851])).
fof(f3993, plain, (spl66_56 | ~ spl66_51 | ~ spl66_139), inference(avatar_split_clause, [], [f3992, f1450, f802, f823])).
fof(f3992, plain, ((e3 = op(e0, e2)) | (~ spl66_51 | ~ spl66_139)), inference(forward_demodulation, [], [f1451, f804])).
fof(f1451, plain, ((e3 = op(e0, op(e0, e3))) | ~ spl66_139), inference(avatar_component_clause, [], [f1450])).
fof(f3991, plain, (spl66_62 | ~ spl66_57 | ~ spl66_141), inference(avatar_split_clause, [], [f3990, f1460, f828, f849])).
fof(f3990, plain, ((op(e0, e0) = e1) | (~ spl66_57 | ~ spl66_141)), inference(forward_demodulation, [], [f1461, f830])).
fof(f3989, plain, (~ spl66_1 | spl66_16 | ~ spl66_143), inference(avatar_contradiction_clause, [], [f3988])).
fof(f3988, plain, ($false | (~ spl66_1 | spl66_16 | ~ spl66_143)), inference(subsumption_resolution, [], [f3987, f654])).
fof(f654, plain, (~ (e3 = op(e3, e0)) | spl66_16), inference(avatar_component_clause, [], [f653])).
fof(f3987, plain, ((e3 = op(e3, e0)) | (~ spl66_1 | ~ spl66_143)), inference(forward_demodulation, [], [f1475, f592])).
fof(f1475, plain, ((e3 = op(e3, op(e3, e3))) | ~ spl66_143), inference(avatar_component_clause, [], [f1474])).
fof(f1474, plain, (spl66_143 <=> (e3 = op(e3, op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl66_143])])).
fof(f3978, plain, (~ spl66_1 | spl66_49 | ~ spl66_158), inference(avatar_contradiction_clause, [], [f3977])).
fof(f3977, plain, ($false | (~ spl66_1 | spl66_49 | ~ spl66_158)), inference(subsumption_resolution, [], [f3976, f795])).
fof(f795, plain, (~ (e0 = op(e0, e3)) | spl66_49), inference(avatar_component_clause, [], [f794])).
fof(f3976, plain, ((e0 = op(e0, e3)) | (~ spl66_1 | ~ spl66_158)), inference(forward_demodulation, [], [f1551, f592])).
fof(f3958, plain, (~ spl66_24 | ~ spl66_28), inference(avatar_split_clause, [], [f3957, f704, f687])).
fof(f687, plain, (spl66_24 <=> (e3 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_24])])).
fof(f3957, plain, (~ (e3 = op(e2, e2)) | ~ spl66_28), inference(forward_demodulation, [], [f274, f706])).
fof(f3955, plain, (~ spl66_17 | ~ spl66_1), inference(avatar_split_clause, [], [f3656, f590, f658])).
fof(f3656, plain, (~ (e0 = op(e2, e3)) | ~ spl66_1), inference(forward_demodulation, [], [f258, f592])).
fof(f3954, plain, (~ spl66_13 | ~ spl66_1), inference(avatar_split_clause, [], [f3655, f590, f641])).
fof(f3655, plain, (~ (e0 = op(e3, e0)) | ~ spl66_1), inference(forward_demodulation, [], [f279, f592])).
fof(f3945, plain, (~ spl66_11 | ~ spl66_14 | ~ spl66_64 | spl66_165), inference(avatar_contradiction_clause, [], [f3944])).
fof(f3944, plain, ($false | (~ spl66_11 | ~ spl66_14 | ~ spl66_64 | spl66_165)), inference(subsumption_resolution, [], [f3943, f634])).
fof(f3943, plain, (~ (e2 = op(e3, e1)) | (~ spl66_14 | ~ spl66_64 | spl66_165)), inference(forward_demodulation, [], [f3942, f647])).
fof(f3942, plain, (~ (e2 = op(e3, op(e3, e0))) | (~ spl66_64 | spl66_165)), inference(forward_demodulation, [], [f1587, f859])).
fof(f859, plain, ((op(e0, e0) = e3) | ~ spl66_64), inference(avatar_component_clause, [], [f857])).
fof(f1587, plain, (~ (e2 = op(op(e0, e0), op(op(e0, e0), e0))) | spl66_165), inference(avatar_component_clause, [], [f1585])).
fof(f1585, plain, (spl66_165 <=> (e2 = op(op(e0, e0), op(op(e0, e0), e0)))), introduced(avatar_definition, [new_symbols(naming, [spl66_165])])).
fof(f3938, plain, (~ spl66_6 | ~ spl66_11 | spl66_145), inference(avatar_split_clause, [], [f3937, f1484, f632, f611])).
fof(f611, plain, (spl66_6 <=> (e1 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_6])])).
fof(f1484, plain, (spl66_145 <=> (e1 = op(e3, op(e3, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl66_145])])).
fof(f3937, plain, (~ (e1 = op(e3, e2)) | (~ spl66_11 | spl66_145)), inference(backward_demodulation, [], [f1486, f634])).
fof(f1486, plain, (~ (e1 = op(e3, op(e3, e1))) | spl66_145), inference(avatar_component_clause, [], [f1484])).
fof(f3933, plain, (~ spl66_6 | ~ spl66_14), inference(avatar_split_clause, [], [f3927, f645, f611])).
fof(f3927, plain, (~ (e1 = op(e3, e2)) | ~ spl66_14), inference(backward_demodulation, [], [f278, f647])).
fof(f3920, plain, (~ spl66_12 | ~ spl66_28), inference(avatar_split_clause, [], [f3917, f704, f636])).
fof(f3917, plain, (~ (e3 = op(e3, e1)) | ~ spl66_28), inference(backward_demodulation, [], [f246, f706])).
fof(f246, plain, ~ (op(e2, e1) = op(e3, e1)), inference(cnf_transformation, [], [f3])).
fof(f3908, plain, (~ spl66_20 | ~ spl66_36), inference(avatar_split_clause, [], [f3906, f738, f670])).
fof(f3906, plain, (~ (e3 = op(e2, e3)) | ~ spl66_36), inference(backward_demodulation, [], [f256, f740])).
fof(f740, plain, ((e3 = op(e1, e3)) | ~ spl66_36), inference(avatar_component_clause, [], [f738])).
fof(f3899, plain, (spl66_37 | ~ spl66_47 | ~ spl66_137), inference(avatar_split_clause, [], [f3896, f1441, f785, f743])).
fof(f3896, plain, ((e0 = op(e1, e2)) | (~ spl66_47 | ~ spl66_137)), inference(backward_demodulation, [], [f1442, f787])).
fof(f3898, plain, (~ spl66_35 | ~ spl66_47), inference(avatar_split_clause, [], [f3894, f785, f734])).
fof(f3894, plain, (~ (e2 = op(e1, e3)) | ~ spl66_47), inference(backward_demodulation, [], [f267, f787])).
fof(f3897, plain, (~ spl66_15 | ~ spl66_47), inference(avatar_split_clause, [], [f3892, f785, f649])).
fof(f3892, plain, (~ (e2 = op(e3, e0)) | ~ spl66_47), inference(backward_demodulation, [], [f239, f787])).
fof(f239, plain, ~ (op(e1, e0) = op(e3, e0)), inference(cnf_transformation, [], [f3])).
fof(f3887, plain, (~ spl66_35 | ~ spl66_51), inference(avatar_split_clause, [], [f3881, f802, f734])).
fof(f3881, plain, (~ (e2 = op(e1, e3)) | ~ spl66_51), inference(backward_demodulation, [], [f253, f804])).
fof(f3879, plain, (~ spl66_50 | ~ spl66_54), inference(avatar_split_clause, [], [f3874, f815, f798])).
fof(f3874, plain, (~ (e1 = op(e0, e3)) | ~ spl66_54), inference(backward_demodulation, [], [f264, f817])).
fof(f3878, plain, (~ spl66_6 | ~ spl66_54), inference(avatar_split_clause, [], [f3873, f815, f611])).
fof(f3873, plain, (~ (e1 = op(e3, e2)) | ~ spl66_54), inference(backward_demodulation, [], [f249, f817])).
fof(f3871, plain, (~ spl66_53 | ~ spl66_57), inference(avatar_split_clause, [], [f3864, f828, f811])).
fof(f3864, plain, (~ (e0 = op(e0, e2)) | ~ spl66_57), inference(backward_demodulation, [], [f262, f830])).
fof(f3870, plain, (~ spl66_25 | ~ spl66_57), inference(avatar_split_clause, [], [f3862, f828, f692])).
fof(f3862, plain, (~ (e0 = op(e2, e1)) | ~ spl66_57), inference(backward_demodulation, [], [f242, f830])).
fof(f3860, plain, (spl66_50 | ~ spl66_1 | ~ spl66_151), inference(avatar_split_clause, [], [f3856, f1516, f590, f798])).
fof(f1516, plain, (spl66_151 <=> (e1 = op(op(e3, e3), e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_151])])).
fof(f3856, plain, ((e1 = op(e0, e3)) | (~ spl66_1 | ~ spl66_151)), inference(forward_demodulation, [], [f1517, f592])).
fof(f1517, plain, ((e1 = op(op(e3, e3), e3)) | ~ spl66_151), inference(avatar_component_clause, [], [f1516])).
fof(f3857, plain, (~ spl66_14 | ~ spl66_64 | spl66_167), inference(avatar_split_clause, [], [f3769, f1595, f857, f645])).
fof(f3769, plain, (~ (e1 = op(e3, e0)) | (~ spl66_64 | spl66_167)), inference(backward_demodulation, [], [f1597, f859])).
fof(f3854, plain, (~ spl66_6 | ~ spl66_15 | ~ spl66_64 | spl66_159), inference(avatar_contradiction_clause, [], [f3853])).
fof(f3853, plain, ($false | (~ spl66_6 | ~ spl66_15 | ~ spl66_64 | spl66_159)), inference(subsumption_resolution, [], [f3852, f613])).
fof(f613, plain, ((e1 = op(e3, e2)) | ~ spl66_6), inference(avatar_component_clause, [], [f611])).
fof(f3852, plain, (~ (e1 = op(e3, e2)) | (~ spl66_15 | ~ spl66_64 | spl66_159)), inference(forward_demodulation, [], [f3851, f651])).
fof(f3851, plain, (~ (e1 = op(e3, op(e3, e0))) | (~ spl66_64 | spl66_159)), inference(forward_demodulation, [], [f1557, f859])).
fof(f1557, plain, (~ (e1 = op(op(e0, e0), op(op(e0, e0), e0))) | spl66_159), inference(avatar_component_clause, [], [f1555])).
fof(f1555, plain, (spl66_159 <=> (e1 = op(op(e0, e0), op(op(e0, e0), e0)))), introduced(avatar_definition, [new_symbols(naming, [spl66_159])])).
fof(f3842, plain, (~ spl66_11 | ~ spl66_15), inference(avatar_split_clause, [], [f3838, f649, f632])).
fof(f3838, plain, (~ (e2 = op(e3, e1)) | ~ spl66_15), inference(backward_demodulation, [], [f277, f651])).
fof(f3832, plain, (spl66_25 | ~ spl66_30 | ~ spl66_132), inference(avatar_split_clause, [], [f3830, f1417, f713, f692])).
fof(f1417, plain, (spl66_132 <=> (e0 = op(e2, op(e2, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl66_132])])).
fof(f3830, plain, ((e0 = op(e2, e1)) | (~ spl66_30 | ~ spl66_132)), inference(backward_demodulation, [], [f1418, f715])).
fof(f1418, plain, ((e0 = op(e2, op(e2, e0))) | ~ spl66_132), inference(avatar_component_clause, [], [f1417])).
fof(f3831, plain, (~ spl66_26 | ~ spl66_30), inference(avatar_split_clause, [], [f3827, f713, f696])).
fof(f696, plain, (spl66_26 <=> (e1 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_26])])).
fof(f3827, plain, (~ (e1 = op(e2, e1)) | ~ spl66_30), inference(backward_demodulation, [], [f271, f715])).
fof(f3817, plain, (~ spl66_34 | ~ spl66_42), inference(avatar_split_clause, [], [f3810, f764, f730])).
fof(f3810, plain, (~ (e1 = op(e1, e3)) | ~ spl66_42), inference(backward_demodulation, [], [f269, f766])).
fof(f3816, plain, (~ spl66_26 | ~ spl66_42), inference(avatar_split_clause, [], [f3807, f764, f696])).
fof(f3807, plain, (~ (e1 = op(e2, e1)) | ~ spl66_42), inference(backward_demodulation, [], [f244, f766])).
fof(f3806, plain, (~ spl66_37 | ~ spl66_45), inference(avatar_split_clause, [], [f3799, f777, f743])).
fof(f3799, plain, (~ (e0 = op(e1, e2)) | ~ spl66_45), inference(backward_demodulation, [], [f266, f779])).
fof(f3789, plain, (~ spl66_37 | ~ spl66_53), inference(avatar_split_clause, [], [f3784, f811, f743])).
fof(f3784, plain, (~ (e0 = op(e1, e2)) | ~ spl66_53), inference(backward_demodulation, [], [f247, f813])).
fof(f3783, plain, (~ spl66_51 | ~ spl66_59), inference(avatar_split_clause, [], [f3779, f836, f802])).
fof(f3779, plain, (~ (e2 = op(e0, e3)) | ~ spl66_59), inference(backward_demodulation, [], [f263, f838])).
fof(f3782, plain, (~ spl66_11 | ~ spl66_59), inference(avatar_split_clause, [], [f3777, f836, f632])).
fof(f3777, plain, (~ (e2 = op(e3, e1)) | ~ spl66_59), inference(backward_demodulation, [], [f243, f838])).
fof(f3774, plain, (~ spl66_15 | ~ spl66_64 | spl66_162), inference(avatar_split_clause, [], [f3768, f1569, f857, f649])).
fof(f3768, plain, (~ (e2 = op(e3, e0)) | (~ spl66_64 | spl66_162)), inference(backward_demodulation, [], [f1571, f859])).
fof(f1571, plain, (~ (e2 = op(op(e0, e0), e0)) | spl66_162), inference(avatar_component_clause, [], [f1569])).
fof(f3772, plain, (~ spl66_56 | ~ spl66_64), inference(avatar_split_clause, [], [f3765, f857, f823])).
fof(f3765, plain, (~ (e3 = op(e0, e2)) | ~ spl66_64), inference(backward_demodulation, [], [f260, f859])).
fof(f3771, plain, (~ spl66_16 | ~ spl66_64), inference(avatar_split_clause, [], [f3763, f857, f653])).
fof(f3763, plain, (~ (e3 = op(e3, e0)) | ~ spl66_64), inference(backward_demodulation, [], [f237, f859])).
fof(f3760, plain, (~ spl66_16 | ~ spl66_1 | spl66_143), inference(avatar_split_clause, [], [f3759, f1474, f590, f653])).
fof(f3759, plain, (~ (e3 = op(e3, e0)) | (~ spl66_1 | spl66_143)), inference(forward_demodulation, [], [f1476, f592])).
fof(f1476, plain, (~ (e3 = op(e3, op(e3, e3))) | spl66_143), inference(avatar_component_clause, [], [f1474])).
fof(f3756, plain, (~ spl66_50 | ~ spl66_1 | spl66_151), inference(avatar_split_clause, [], [f3649, f1516, f590, f798])).
fof(f3649, plain, (~ (e1 = op(e0, e3)) | (~ spl66_1 | spl66_151)), inference(backward_demodulation, [], [f1518, f592])).
fof(f1518, plain, (~ (e1 = op(op(e3, e3), e3)) | spl66_151), inference(avatar_component_clause, [], [f1516])).
fof(f3738, plain, (~ spl66_8 | ~ spl66_16), inference(avatar_split_clause, [], [f3736, f653, f619])).
fof(f3736, plain, (~ (e3 = op(e3, e2)) | ~ spl66_16), inference(backward_demodulation, [], [f278, f655])).
fof(f655, plain, ((e3 = op(e3, e0)) | ~ spl66_16), inference(avatar_component_clause, [], [f653])).
fof(f3733, plain, (~ spl66_18 | ~ spl66_26), inference(avatar_split_clause, [], [f3728, f696, f662])).
fof(f3728, plain, (~ (e1 = op(e2, e3)) | ~ spl66_26), inference(backward_demodulation, [], [f275, f698])).
fof(f698, plain, ((e1 = op(e2, e1)) | ~ spl66_26), inference(avatar_component_clause, [], [f696])).
fof(f3732, plain, (~ spl66_10 | ~ spl66_26), inference(avatar_split_clause, [], [f3727, f696, f628])).
fof(f3727, plain, (~ (e1 = op(e3, e1)) | ~ spl66_26), inference(backward_demodulation, [], [f246, f698])).
fof(f3698, plain, (~ spl66_51 | ~ spl66_56 | spl66_139), inference(avatar_contradiction_clause, [], [f3697])).
fof(f3697, plain, ($false | (~ spl66_51 | ~ spl66_56 | spl66_139)), inference(subsumption_resolution, [], [f3693, f825])).
fof(f3693, plain, (~ (e3 = op(e0, e2)) | (~ spl66_51 | spl66_139)), inference(backward_demodulation, [], [f1452, f804])).
fof(f3688, plain, (~ spl66_52 | ~ spl66_56), inference(avatar_split_clause, [], [f3685, f823, f806])).
fof(f3685, plain, (~ (e3 = op(e0, e3)) | ~ spl66_56), inference(backward_demodulation, [], [f264, f825])).
fof(f3664, plain, (spl66_11 | ~ spl66_44 | ~ spl66_155), inference(avatar_split_clause, [], [f3663, f1535, f772, f632])).
fof(f772, plain, (spl66_44 <=> (e3 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_44])])).
fof(f3663, plain, ((e2 = op(e3, e1)) | (~ spl66_44 | ~ spl66_155)), inference(forward_demodulation, [], [f1536, f774])).
fof(f774, plain, ((e3 = op(e1, e1)) | ~ spl66_44), inference(avatar_component_clause, [], [f772])).
fof(f1536, plain, ((e2 = op(op(e1, e1), e1)) | ~ spl66_155), inference(avatar_component_clause, [], [f1535])).
fof(f3642, plain, (~ spl66_2 | ~ spl66_10), inference(avatar_split_clause, [], [f3638, f628, f594])).
fof(f594, plain, (spl66_2 <=> (e1 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_2])])).
fof(f3638, plain, (~ (e1 = op(e3, e3)) | ~ spl66_10), inference(backward_demodulation, [], [f281, f630])).
fof(f630, plain, ((e1 = op(e3, e1)) | ~ spl66_10), inference(avatar_component_clause, [], [f628])).
fof(f3627, plain, (~ spl66_2 | ~ spl66_18), inference(avatar_split_clause, [], [f3623, f662, f594])).
fof(f3623, plain, (~ (e1 = op(e3, e3)) | ~ spl66_18), inference(backward_demodulation, [], [f258, f664])).
fof(f664, plain, ((e1 = op(e2, e3)) | ~ spl66_18), inference(avatar_component_clause, [], [f662])).
fof(f3622, plain, (~ spl66_20 | ~ spl66_32), inference(avatar_split_clause, [], [f3620, f721, f670])).
fof(f3620, plain, (~ (e3 = op(e2, e3)) | ~ spl66_32), inference(backward_demodulation, [], [f273, f723])).
fof(f3621, plain, (~ spl66_16 | ~ spl66_32), inference(avatar_split_clause, [], [f3619, f721, f653])).
fof(f3619, plain, (~ (e3 = op(e3, e0)) | ~ spl66_32), inference(backward_demodulation, [], [f240, f723])).
fof(f240, plain, ~ (op(e2, e0) = op(e3, e0)), inference(cnf_transformation, [], [f3])).
fof(f3612, plain, (~ spl66_45 | spl66_137), inference(avatar_contradiction_clause, [], [f3611])).
fof(f3611, plain, ($false | (~ spl66_45 | spl66_137)), inference(subsumption_resolution, [], [f3608, f779])).
fof(f3608, plain, (~ (e0 = op(e1, e0)) | (~ spl66_45 | spl66_137)), inference(backward_demodulation, [], [f1443, f779])).
fof(f1443, plain, (~ (e0 = op(e1, op(e1, e0))) | spl66_137), inference(avatar_component_clause, [], [f1441])).
fof(f3610, plain, (~ spl66_33 | ~ spl66_45), inference(avatar_split_clause, [], [f3604, f777, f726])).
fof(f3604, plain, (~ (e0 = op(e1, e3)) | ~ spl66_45), inference(backward_demodulation, [], [f267, f779])).
fof(f3601, plain, (~ spl66_20 | ~ spl66_52), inference(avatar_split_clause, [], [f3598, f806, f670])).
fof(f3598, plain, (~ (e3 = op(e2, e3)) | ~ spl66_52), inference(backward_demodulation, [], [f254, f808])).
fof(f808, plain, ((e3 = op(e0, e3)) | ~ spl66_52), inference(avatar_component_clause, [], [f806])).
fof(f3583, plain, (~ spl66_47 | ~ spl66_62 | spl66_162), inference(avatar_split_clause, [], [f3579, f1569, f849, f785])).
fof(f3579, plain, (~ (e2 = op(e1, e0)) | (~ spl66_62 | spl66_162)), inference(backward_demodulation, [], [f1571, f851])).
fof(f3582, plain, (~ spl66_58 | ~ spl66_62), inference(avatar_split_clause, [], [f3573, f849, f832])).
fof(f832, plain, (spl66_58 <=> (e1 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_58])])).
fof(f3573, plain, (~ (e1 = op(e0, e1)) | ~ spl66_62), inference(backward_demodulation, [], [f259, f851])).
fof(f3568, plain, (~ spl66_11 | ~ spl66_44 | spl66_155), inference(avatar_split_clause, [], [f3567, f1535, f772, f632])).
fof(f3567, plain, (~ (e2 = op(e3, e1)) | (~ spl66_44 | spl66_155)), inference(forward_demodulation, [], [f1537, f774])).
fof(f3558, plain, (spl66_30 | ~ spl66_25 | ~ spl66_131), inference(avatar_split_clause, [], [f3527, f1412, f692, f713])).
fof(f3527, plain, ((e1 = op(e2, e0)) | (~ spl66_25 | ~ spl66_131)), inference(backward_demodulation, [], [f1413, f694])).
fof(f3557, plain, (~ spl66_5 | spl66_15 | ~ spl66_144), inference(avatar_contradiction_clause, [], [f3556])).
fof(f3556, plain, ($false | (~ spl66_5 | spl66_15 | ~ spl66_144)), inference(subsumption_resolution, [], [f3555, f650])).
fof(f650, plain, (~ (e2 = op(e3, e0)) | spl66_15), inference(avatar_component_clause, [], [f649])).
fof(f3555, plain, ((e2 = op(e3, e0)) | (~ spl66_5 | ~ spl66_144)), inference(backward_demodulation, [], [f1480, f609])).
fof(f609, plain, ((e0 = op(e3, e2)) | ~ spl66_5), inference(avatar_component_clause, [], [f607])).
fof(f1480, plain, ((e2 = op(e3, op(e3, e2))) | ~ spl66_144), inference(avatar_component_clause, [], [f1479])).
fof(f1479, plain, (spl66_144 <=> (e2 = op(e3, op(e3, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl66_144])])).
fof(f3551, plain, (~ spl66_5 | ~ spl66_11 | ~ spl66_44 | spl66_152), inference(avatar_split_clause, [], [f3548, f1521, f772, f632, f607])).
fof(f1521, plain, (spl66_152 <=> (e0 = op(op(e1, e1), op(op(e1, e1), e1)))), introduced(avatar_definition, [new_symbols(naming, [spl66_152])])).
fof(f3548, plain, (~ (e0 = op(e3, e2)) | (~ spl66_11 | ~ spl66_44 | spl66_152)), inference(backward_demodulation, [], [f3493, f634])).
fof(f3493, plain, (~ (e0 = op(e3, op(e3, e1))) | (~ spl66_44 | spl66_152)), inference(backward_demodulation, [], [f1523, f774])).
fof(f1523, plain, (~ (e0 = op(op(e1, e1), op(op(e1, e1), e1))) | spl66_152), inference(avatar_component_clause, [], [f1521])).
fof(f3540, plain, (~ spl66_7 | ~ spl66_23), inference(avatar_split_clause, [], [f3530, f683, f615])).
fof(f3530, plain, (~ (e2 = op(e3, e2)) | ~ spl66_23), inference(backward_demodulation, [], [f252, f685])).
fof(f3528, plain, (~ spl66_9 | ~ spl66_25), inference(avatar_split_clause, [], [f3521, f692, f624])).
fof(f3521, plain, (~ (e0 = op(e3, e1)) | ~ spl66_25), inference(backward_demodulation, [], [f246, f694])).
fof(f3514, plain, (~ spl66_2 | ~ spl66_33 | ~ spl66_47 | spl66_163), inference(avatar_contradiction_clause, [], [f3513])).
fof(f3513, plain, ($false | (~ spl66_2 | ~ spl66_33 | ~ spl66_47 | spl66_163)), inference(subsumption_resolution, [], [f3509, f787])).
fof(f3509, plain, (~ (e2 = op(e1, e0)) | (~ spl66_2 | ~ spl66_33 | spl66_163)), inference(backward_demodulation, [], [f3436, f728])).
fof(f3436, plain, (~ (e2 = op(e1, op(e1, e3))) | (~ spl66_2 | spl66_163)), inference(forward_demodulation, [], [f1576, f596])).
fof(f596, plain, ((e1 = op(e3, e3)) | ~ spl66_2), inference(avatar_component_clause, [], [f594])).
fof(f1576, plain, (~ (e2 = op(op(e3, e3), op(op(e3, e3), e3))) | spl66_163), inference(avatar_component_clause, [], [f1574])).
fof(f1574, plain, (spl66_163 <=> (e2 = op(op(e3, e3), op(op(e3, e3), e3)))), introduced(avatar_definition, [new_symbols(naming, [spl66_163])])).
fof(f3497, plain, (~ spl66_9 | ~ spl66_44 | spl66_166), inference(avatar_split_clause, [], [f3494, f1590, f772, f624])).
fof(f3494, plain, (~ (e0 = op(e3, e1)) | (~ spl66_44 | spl66_166)), inference(backward_demodulation, [], [f1592, f774])).
fof(f3496, plain, (~ spl66_36 | ~ spl66_44), inference(avatar_split_clause, [], [f3491, f772, f738])).
fof(f3491, plain, (~ (e3 = op(e1, e3)) | ~ spl66_44), inference(backward_demodulation, [], [f269, f774])).
fof(f3487, plain, (~ spl66_37 | ~ spl66_47 | spl66_137), inference(avatar_split_clause, [], [f3484, f1441, f785, f743])).
fof(f3484, plain, (~ (e0 = op(e1, e2)) | (~ spl66_47 | spl66_137)), inference(backward_demodulation, [], [f1443, f787])).
fof(f3486, plain, (~ spl66_43 | ~ spl66_47), inference(avatar_split_clause, [], [f3480, f785, f768])).
fof(f3480, plain, (~ (e2 = op(e1, e1)) | ~ spl66_47), inference(backward_demodulation, [], [f265, f787])).
fof(f3485, plain, (~ spl66_31 | ~ spl66_47), inference(avatar_split_clause, [], [f3479, f785, f717])).
fof(f3479, plain, (~ (e2 = op(e2, e0)) | ~ spl66_47), inference(backward_demodulation, [], [f238, f787])).
fof(f238, plain, ~ (op(e1, e0) = op(e2, e0)), inference(cnf_transformation, [], [f3])).
fof(f3478, plain, (~ spl66_24 | ~ spl66_56), inference(avatar_split_clause, [], [f3476, f823, f687])).
fof(f3476, plain, (~ (e3 = op(e2, e2)) | ~ spl66_56), inference(backward_demodulation, [], [f248, f825])).
fof(f3474, plain, (~ spl66_54 | ~ spl66_58), inference(avatar_split_clause, [], [f3469, f832, f815])).
fof(f3469, plain, (~ (e1 = op(e0, e2)) | ~ spl66_58), inference(backward_demodulation, [], [f262, f834])).
fof(f834, plain, ((e1 = op(e0, e1)) | ~ spl66_58), inference(avatar_component_clause, [], [f832])).
fof(f3473, plain, (~ spl66_26 | ~ spl66_58), inference(avatar_split_clause, [], [f3467, f832, f696])).
fof(f3467, plain, (~ (e1 = op(e2, e1)) | ~ spl66_58), inference(backward_demodulation, [], [f242, f834])).
fof(f3459, plain, (~ spl66_48 | ~ spl66_16), inference(avatar_split_clause, [], [f3458, f653, f789])).
fof(f3458, plain, (~ (e3 = op(e1, e0)) | ~ spl66_16), inference(forward_demodulation, [], [f239, f655])).
fof(f3440, plain, (~ spl66_7 | spl66_144), inference(avatar_contradiction_clause, [], [f3439])).
fof(f3439, plain, ($false | (~ spl66_7 | spl66_144)), inference(subsumption_resolution, [], [f3438, f617])).
fof(f617, plain, ((e2 = op(e3, e2)) | ~ spl66_7), inference(avatar_component_clause, [], [f615])).
fof(f3438, plain, (~ (e2 = op(e3, e2)) | (~ spl66_7 | spl66_144)), inference(forward_demodulation, [], [f1481, f617])).
fof(f1481, plain, (~ (e2 = op(e3, op(e3, e2))) | spl66_144), inference(avatar_component_clause, [], [f1479])).
fof(f3424, plain, (~ spl66_8 | ~ spl66_24), inference(avatar_split_clause, [], [f3419, f687, f619])).
fof(f3419, plain, (~ (e3 = op(e3, e2)) | ~ spl66_24), inference(backward_demodulation, [], [f252, f689])).
fof(f689, plain, ((e3 = op(e2, e2)) | ~ spl66_24), inference(avatar_component_clause, [], [f687])).
fof(f3407, plain, (~ spl66_35 | ~ spl66_43), inference(avatar_split_clause, [], [f3403, f768, f734])).
fof(f3403, plain, (~ (e2 = op(e1, e3)) | ~ spl66_43), inference(backward_demodulation, [], [f269, f770])).
fof(f3397, plain, (~ spl66_52 | ~ spl66_60), inference(avatar_split_clause, [], [f3394, f840, f806])).
fof(f3394, plain, (~ (e3 = op(e0, e3)) | ~ spl66_60), inference(backward_demodulation, [], [f263, f842])).
fof(f3396, plain, (~ spl66_44 | ~ spl66_60), inference(avatar_split_clause, [], [f3393, f840, f772])).
fof(f3393, plain, (~ (e3 = op(e1, e1)) | ~ spl66_60), inference(backward_demodulation, [], [f241, f842])).
fof(f3375, plain, (~ spl66_35 | ~ spl66_2 | spl66_148), inference(avatar_split_clause, [], [f3362, f1502, f594, f734])).
fof(f3362, plain, (~ (e2 = op(e1, e3)) | (~ spl66_2 | spl66_148)), inference(backward_demodulation, [], [f1504, f596])).
fof(f1504, plain, (~ (e2 = op(op(e3, e3), e3)) | spl66_148), inference(avatar_component_clause, [], [f1502])).
fof(f3372, plain, (~ spl66_32 | ~ spl66_17 | spl66_129), inference(avatar_split_clause, [], [f3332, f1402, f658, f721])).
fof(f1402, plain, (spl66_129 <=> (e3 = op(e2, op(e2, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl66_129])])).
fof(f3332, plain, (~ (e3 = op(e2, e0)) | (~ spl66_17 | spl66_129)), inference(backward_demodulation, [], [f1404, f660])).
fof(f660, plain, ((e0 = op(e2, e3)) | ~ spl66_17), inference(avatar_component_clause, [], [f658])).
fof(f1404, plain, (~ (e3 = op(e2, op(e2, e3))) | spl66_129), inference(avatar_component_clause, [], [f1402])).
fof(f3340, plain, (~ spl66_3 | ~ spl66_15), inference(avatar_split_clause, [], [f3337, f649, f598])).
fof(f3337, plain, (~ (e2 = op(e3, e3)) | ~ spl66_15), inference(backward_demodulation, [], [f279, f651])).
fof(f3324, plain, (~ spl66_23 | spl66_130), inference(avatar_contradiction_clause, [], [f3323])).
fof(f3323, plain, ($false | (~ spl66_23 | spl66_130)), inference(subsumption_resolution, [], [f3319, f685])).
fof(f3319, plain, (~ (e2 = op(e2, e2)) | (~ spl66_23 | spl66_130)), inference(backward_demodulation, [], [f1409, f685])).
fof(f1409, plain, (~ (e2 = op(e2, op(e2, e2))) | spl66_130), inference(avatar_component_clause, [], [f1407])).
fof(f1407, plain, (spl66_130 <=> (e2 = op(e2, op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl66_130])])).
fof(f3315, plain, (~ spl66_26 | spl66_131), inference(avatar_contradiction_clause, [], [f3314])).
fof(f3314, plain, ($false | (~ spl66_26 | spl66_131)), inference(subsumption_resolution, [], [f3311, f698])).
fof(f3311, plain, (~ (e1 = op(e2, e1)) | (~ spl66_26 | spl66_131)), inference(backward_demodulation, [], [f1414, f698])).
fof(f1414, plain, (~ (e1 = op(e2, op(e2, e1))) | spl66_131), inference(avatar_component_clause, [], [f1412])).
fof(f3305, plain, (~ spl66_17 | ~ spl66_32 | spl66_132), inference(avatar_split_clause, [], [f3302, f1417, f721, f658])).
fof(f3302, plain, (~ (e0 = op(e2, e3)) | (~ spl66_32 | spl66_132)), inference(backward_demodulation, [], [f1419, f723])).
fof(f1419, plain, (~ (e0 = op(e2, op(e2, e0))) | spl66_132), inference(avatar_component_clause, [], [f1417])).
fof(f3304, plain, (~ spl66_24 | ~ spl66_32), inference(avatar_split_clause, [], [f3300, f721, f687])).
fof(f3300, plain, (~ (e3 = op(e2, e2)) | ~ spl66_32), inference(backward_demodulation, [], [f272, f723])).
fof(f3294, plain, (~ spl66_3 | ~ spl66_35), inference(avatar_split_clause, [], [f3290, f734, f598])).
fof(f3290, plain, (~ (e2 = op(e3, e3)) | ~ spl66_35), inference(backward_demodulation, [], [f257, f736])).
fof(f3268, plain, (~ spl66_52 | spl66_139), inference(avatar_contradiction_clause, [], [f3267])).
fof(f3267, plain, ($false | (~ spl66_52 | spl66_139)), inference(subsumption_resolution, [], [f3265, f808])).
fof(f3265, plain, (~ (e3 = op(e0, e3)) | (~ spl66_52 | spl66_139)), inference(backward_demodulation, [], [f1452, f808])).
fof(f3266, plain, (~ spl66_36 | ~ spl66_52), inference(avatar_split_clause, [], [f3262, f806, f738])).
fof(f3262, plain, (~ (e3 = op(e1, e3)) | ~ spl66_52), inference(backward_demodulation, [], [f253, f808])).
fof(f3261, plain, (~ spl66_27 | ~ spl66_59), inference(avatar_split_clause, [], [f3256, f836, f700])).
fof(f3256, plain, (~ (e2 = op(e2, e1)) | ~ spl66_59), inference(backward_demodulation, [], [f242, f838])).
fof(f3248, plain, (~ spl66_49 | ~ spl66_61), inference(avatar_split_clause, [], [f3238, f845, f794])).
fof(f3238, plain, (~ (e0 = op(e0, e3)) | ~ spl66_61), inference(backward_demodulation, [], [f261, f847])).
fof(f3247, plain, (~ spl66_29 | ~ spl66_61), inference(avatar_split_clause, [], [f3235, f845, f709])).
fof(f709, plain, (spl66_29 <=> (e0 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_29])])).
fof(f3235, plain, (~ (e0 = op(e2, e0)) | ~ spl66_61), inference(backward_demodulation, [], [f236, f847])).
fof(f3234, plain, (~ spl66_41 | ~ spl66_46 | spl66_137), inference(avatar_split_clause, [], [f3233, f1441, f781, f760])).
fof(f3233, plain, (~ (e0 = op(e1, e1)) | (~ spl66_46 | spl66_137)), inference(forward_demodulation, [], [f1443, f783])).
fof(f3227, plain, (~ spl66_42 | ~ spl66_46), inference(avatar_split_clause, [], [f3226, f781, f764])).
fof(f3226, plain, (~ (e1 = op(e1, e1)) | ~ spl66_46), inference(forward_demodulation, [], [f265, f783])).
fof(f3194, plain, (~ spl66_29 | spl66_132), inference(avatar_contradiction_clause, [], [f3193])).
fof(f3193, plain, ($false | (~ spl66_29 | spl66_132)), inference(subsumption_resolution, [], [f3189, f711])).
fof(f711, plain, ((e0 = op(e2, e0)) | ~ spl66_29), inference(avatar_component_clause, [], [f709])).
fof(f3189, plain, (~ (e0 = op(e2, e0)) | (~ spl66_29 | spl66_132)), inference(backward_demodulation, [], [f1419, f711])).
fof(f3192, plain, (~ spl66_17 | ~ spl66_29), inference(avatar_split_clause, [], [f3185, f709, f658])).
fof(f3185, plain, (~ (e0 = op(e2, e3)) | ~ spl66_29), inference(backward_demodulation, [], [f273, f711])).
fof(f3181, plain, (~ spl66_41 | ~ spl66_42), inference(avatar_contradiction_clause, [], [f3180])).
fof(f3180, plain, ($false | (~ spl66_41 | ~ spl66_42)), inference(subsumption_resolution, [], [f3179, f283])).
fof(f3179, plain, ((e0 = e1) | (~ spl66_41 | ~ spl66_42)), inference(backward_demodulation, [], [f766, f762])).
fof(f3177, plain, (~ spl66_30 | ~ spl66_46), inference(avatar_split_clause, [], [f3172, f781, f713])).
fof(f3172, plain, (~ (e1 = op(e2, e0)) | ~ spl66_46), inference(backward_demodulation, [], [f238, f783])).
fof(f3154, plain, (spl66_53 | ~ spl66_63 | ~ spl66_142), inference(avatar_contradiction_clause, [], [f3153])).
fof(f3153, plain, ($false | (spl66_53 | ~ spl66_63 | ~ spl66_142)), inference(subsumption_resolution, [], [f3146, f812])).
fof(f812, plain, (~ (e0 = op(e0, e2)) | spl66_53), inference(avatar_component_clause, [], [f811])).
fof(f3146, plain, ((e0 = op(e0, e2)) | (~ spl66_63 | ~ spl66_142)), inference(backward_demodulation, [], [f1466, f855])).
fof(f1466, plain, ((e0 = op(e0, op(e0, e0))) | ~ spl66_142), inference(avatar_component_clause, [], [f1465])).
fof(f1465, plain, (spl66_142 <=> (e0 = op(e0, op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl66_142])])).
fof(f3152, plain, (~ spl66_55 | ~ spl66_63), inference(avatar_split_clause, [], [f3143, f853, f819])).
fof(f3143, plain, (~ (e2 = op(e0, e2)) | ~ spl66_63), inference(backward_demodulation, [], [f260, f855])).
fof(f3139, plain, (spl66_18 | ~ spl66_3 | ~ spl66_151), inference(avatar_split_clause, [], [f3138, f1516, f598, f662])).
fof(f3138, plain, ((e1 = op(e2, e3)) | (~ spl66_3 | ~ spl66_151)), inference(forward_demodulation, [], [f1517, f600])).
fof(f600, plain, ((e2 = op(e3, e3)) | ~ spl66_3), inference(avatar_component_clause, [], [f598])).
fof(f3137, plain, (~ spl66_17 | ~ spl66_3 | spl66_158), inference(avatar_split_clause, [], [f3136, f1550, f598, f658])).
fof(f3136, plain, (~ (e0 = op(e2, e3)) | (~ spl66_3 | spl66_158)), inference(forward_demodulation, [], [f1552, f600])).
fof(f1552, plain, (~ (e0 = op(op(e3, e3), e3)) | spl66_158), inference(avatar_component_clause, [], [f1550])).
fof(f3134, plain, (~ spl66_51 | ~ spl66_3), inference(avatar_split_clause, [], [f3133, f598, f802])).
fof(f3133, plain, (~ (e2 = op(e0, e3)) | ~ spl66_3), inference(forward_demodulation, [], [f255, f600])).
fof(f3121, plain, (~ spl66_31 | ~ spl66_27), inference(avatar_split_clause, [], [f3120, f700, f717])).
fof(f3120, plain, (~ (e2 = op(e2, e0)) | ~ spl66_27), inference(forward_demodulation, [], [f271, f702])).
fof(f3119, plain, (~ spl66_19 | ~ spl66_24 | spl66_130), inference(avatar_split_clause, [], [f3013, f1407, f687, f666])).
fof(f3013, plain, (~ (e2 = op(e2, e3)) | (~ spl66_24 | spl66_130)), inference(forward_demodulation, [], [f1409, f689])).
fof(f3118, plain, (~ spl66_20 | ~ spl66_24), inference(avatar_split_clause, [], [f3117, f687, f670])).
fof(f3117, plain, (~ (e3 = op(e2, e3)) | ~ spl66_24), inference(forward_demodulation, [], [f276, f689])).
fof(f3116, plain, (~ spl66_19 | ~ spl66_27), inference(avatar_split_clause, [], [f3115, f700, f666])).
fof(f3115, plain, (~ (e2 = op(e2, e3)) | ~ spl66_27), inference(forward_demodulation, [], [f275, f702])).
fof(f3110, plain, (~ spl66_12 | ~ spl66_60), inference(avatar_split_clause, [], [f3109, f840, f636])).
fof(f3109, plain, (~ (e3 = op(e3, e1)) | ~ spl66_60), inference(forward_demodulation, [], [f243, f842])).
fof(f3100, plain, (~ spl66_22 | ~ spl66_27 | spl66_131), inference(avatar_split_clause, [], [f3098, f1412, f700, f679])).
fof(f3098, plain, (~ (e1 = op(e2, e2)) | (~ spl66_27 | spl66_131)), inference(backward_demodulation, [], [f1414, f702])).
fof(f3086, plain, (~ spl66_42 | spl66_136), inference(avatar_contradiction_clause, [], [f3085])).
fof(f3085, plain, ($false | (~ spl66_42 | spl66_136)), inference(subsumption_resolution, [], [f3080, f766])).
fof(f3080, plain, (~ (e1 = op(e1, e1)) | (~ spl66_42 | spl66_136)), inference(backward_demodulation, [], [f1438, f766])).
fof(f3067, plain, (~ spl66_9 | spl66_14 | ~ spl66_145), inference(avatar_contradiction_clause, [], [f3066])).
fof(f3066, plain, ($false | (~ spl66_9 | spl66_14 | ~ spl66_145)), inference(subsumption_resolution, [], [f3065, f646])).
fof(f646, plain, (~ (e1 = op(e3, e0)) | spl66_14), inference(avatar_component_clause, [], [f645])).
fof(f3065, plain, ((e1 = op(e3, e0)) | (~ spl66_9 | ~ spl66_145)), inference(forward_demodulation, [], [f1485, f626])).
fof(f626, plain, ((e0 = op(e3, e1)) | ~ spl66_9), inference(avatar_component_clause, [], [f624])).
fof(f1485, plain, ((e1 = op(e3, op(e3, e1))) | ~ spl66_145), inference(avatar_component_clause, [], [f1484])).
fof(f3034, plain, (~ spl66_30 | ~ spl66_3 | ~ spl66_17 | spl66_156), inference(avatar_split_clause, [], [f2981, f1540, f658, f598, f713])).
fof(f1540, plain, (spl66_156 <=> (e1 = op(op(e3, e3), op(op(e3, e3), e3)))), introduced(avatar_definition, [new_symbols(naming, [spl66_156])])).
fof(f2981, plain, (~ (e1 = op(e2, e0)) | (~ spl66_3 | ~ spl66_17 | spl66_156)), inference(backward_demodulation, [], [f2904, f660])).
fof(f2904, plain, (~ (e1 = op(e2, op(e2, e3))) | (~ spl66_3 | spl66_156)), inference(backward_demodulation, [], [f1542, f600])).
fof(f1542, plain, (~ (e1 = op(op(e3, e3), op(op(e3, e3), e3))) | spl66_156), inference(avatar_component_clause, [], [f1540])).
fof(f3019, plain, (~ spl66_22 | ~ spl66_6), inference(avatar_split_clause, [], [f3018, f611, f679])).
fof(f3018, plain, (~ (e1 = op(e2, e2)) | ~ spl66_6), inference(forward_demodulation, [], [f252, f613])).
fof(f2998, plain, (~ spl66_9 | ~ spl66_16 | ~ spl66_44 | spl66_169), inference(avatar_contradiction_clause, [], [f2997])).
fof(f2997, plain, ($false | (~ spl66_9 | ~ spl66_16 | ~ spl66_44 | spl66_169)), inference(subsumption_resolution, [], [f2995, f655])).
fof(f2995, plain, (~ (e3 = op(e3, e0)) | (~ spl66_9 | ~ spl66_44 | spl66_169)), inference(backward_demodulation, [], [f2732, f626])).
fof(f2732, plain, (~ (e3 = op(e3, op(e3, e1))) | (~ spl66_44 | spl66_169)), inference(backward_demodulation, [], [f1607, f774])).
fof(f1607, plain, (~ (e3 = op(op(e1, e1), op(op(e1, e1), e1))) | spl66_169), inference(avatar_component_clause, [], [f1605])).
fof(f1605, plain, (spl66_169 <=> (e3 = op(op(e1, e1), op(op(e1, e1), e1)))), introduced(avatar_definition, [new_symbols(naming, [spl66_169])])).
fof(f2975, plain, (~ spl66_22 | ~ spl66_24), inference(avatar_contradiction_clause, [], [f2974])).
fof(f2974, plain, ($false | (~ spl66_22 | ~ spl66_24)), inference(subsumption_resolution, [], [f2973, f287])).
fof(f287, plain, ~ (e1 = e3), inference(cnf_transformation, [], [f4])).
fof(f2973, plain, ((e1 = e3) | (~ spl66_22 | ~ spl66_24)), inference(forward_demodulation, [], [f689, f681])).
fof(f2967, plain, (~ spl66_36 | spl66_134), inference(avatar_contradiction_clause, [], [f2966])).
fof(f2966, plain, ($false | (~ spl66_36 | spl66_134)), inference(subsumption_resolution, [], [f2964, f740])).
fof(f2964, plain, (~ (e3 = op(e1, e3)) | (~ spl66_36 | spl66_134)), inference(backward_demodulation, [], [f1428, f740])).
fof(f1428, plain, (~ (e3 = op(e1, op(e1, e3))) | spl66_134), inference(avatar_component_clause, [], [f1426])).
fof(f1426, plain, (spl66_134 <=> (e3 = op(e1, op(e1, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl66_134])])).
fof(f2951, plain, (~ spl66_59 | ~ spl66_60), inference(avatar_contradiction_clause, [], [f2950])).
fof(f2950, plain, ($false | (~ spl66_59 | ~ spl66_60)), inference(subsumption_resolution, [], [f2949, f288])).
fof(f288, plain, ~ (e2 = e3), inference(cnf_transformation, [], [f4])).
fof(f2949, plain, ((e2 = e3) | (~ spl66_59 | ~ spl66_60)), inference(forward_demodulation, [], [f842, f838])).
fof(f2932, plain, (~ spl66_54 | ~ spl66_22), inference(avatar_split_clause, [], [f2551, f679, f815])).
fof(f2551, plain, (~ (e1 = op(e0, e2)) | ~ spl66_22), inference(forward_demodulation, [], [f248, f681])).
fof(f2928, plain, (~ spl66_26 | ~ spl66_22), inference(avatar_split_clause, [], [f2557, f679, f696])).
fof(f2557, plain, (~ (e1 = op(e2, e1)) | ~ spl66_22), inference(forward_demodulation, [], [f274, f681])).
fof(f2923, plain, (~ spl66_7 | ~ spl66_3), inference(avatar_split_clause, [], [f2922, f598, f615])).
fof(f2922, plain, (~ (e2 = op(e3, e2)) | ~ spl66_3), inference(forward_demodulation, [], [f282, f600])).
fof(f2921, plain, (~ spl66_8 | ~ spl66_3 | spl66_143), inference(avatar_split_clause, [], [f2901, f1474, f598, f619])).
fof(f2901, plain, (~ (e3 = op(e3, e2)) | (~ spl66_3 | spl66_143)), inference(backward_demodulation, [], [f1476, f600])).
fof(f2914, plain, (~ spl66_61 | spl66_142), inference(avatar_contradiction_clause, [], [f2913])).
fof(f2913, plain, ($false | (~ spl66_61 | spl66_142)), inference(subsumption_resolution, [], [f2912, f847])).
fof(f2912, plain, (~ (e0 = op(e0, e0)) | (~ spl66_61 | spl66_142)), inference(forward_demodulation, [], [f1467, f847])).
fof(f1467, plain, (~ (e0 = op(e0, op(e0, e0))) | spl66_142), inference(avatar_component_clause, [], [f1465])).
fof(f2899, plain, (~ spl66_1 | ~ spl66_5), inference(avatar_split_clause, [], [f2897, f607, f590])).
fof(f2897, plain, (~ (e0 = op(e3, e3)) | ~ spl66_5), inference(backward_demodulation, [], [f282, f609])).
fof(f2878, plain, (~ spl66_19 | ~ spl66_31), inference(avatar_split_clause, [], [f2876, f717, f666])).
fof(f2876, plain, (~ (e2 = op(e2, e3)) | ~ spl66_31), inference(backward_demodulation, [], [f273, f719])).
fof(f2873, plain, (~ spl66_1 | ~ spl66_33), inference(avatar_split_clause, [], [f2868, f726, f590])).
fof(f2868, plain, (~ (e0 = op(e3, e3)) | ~ spl66_33), inference(backward_demodulation, [], [f257, f728])).
fof(f2864, plain, (~ spl66_39 | spl66_135), inference(avatar_contradiction_clause, [], [f2863])).
fof(f2863, plain, ($false | (~ spl66_39 | spl66_135)), inference(subsumption_resolution, [], [f2859, f753])).
fof(f2859, plain, (~ (e2 = op(e1, e2)) | (~ spl66_39 | spl66_135)), inference(backward_demodulation, [], [f1433, f753])).
fof(f1433, plain, (~ (e2 = op(e1, op(e1, e2))) | spl66_135), inference(avatar_component_clause, [], [f1431])).
fof(f1431, plain, (spl66_135 <=> (e2 = op(e1, op(e1, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl66_135])])).
fof(f2804, plain, (~ spl66_34 | ~ spl66_44 | spl66_136), inference(avatar_split_clause, [], [f2729, f1436, f772, f730])).
fof(f2729, plain, (~ (e1 = op(e1, e3)) | (~ spl66_44 | spl66_136)), inference(backward_demodulation, [], [f1438, f774])).
fof(f2801, plain, (spl66_1 | ~ spl66_16 | ~ spl66_146), inference(avatar_split_clause, [], [f2800, f1489, f653, f590])).
fof(f1489, plain, (spl66_146 <=> (e0 = op(e3, op(e3, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl66_146])])).
fof(f2800, plain, ((e0 = op(e3, e3)) | (~ spl66_16 | ~ spl66_146)), inference(forward_demodulation, [], [f1490, f655])).
fof(f1490, plain, ((e0 = op(e3, op(e3, e0))) | ~ spl66_146), inference(avatar_component_clause, [], [f1489])).
fof(f2789, plain, (~ spl66_1 | ~ spl66_49 | spl66_158), inference(avatar_contradiction_clause, [], [f2788])).
fof(f2788, plain, ($false | (~ spl66_1 | ~ spl66_49 | spl66_158)), inference(subsumption_resolution, [], [f2784, f796])).
fof(f2784, plain, (~ (e0 = op(e0, e3)) | (~ spl66_1 | spl66_158)), inference(backward_demodulation, [], [f1552, f592])).
fof(f2775, plain, (~ spl66_1 | ~ spl66_9), inference(avatar_split_clause, [], [f2768, f624, f590])).
fof(f2768, plain, (~ (e0 = op(e3, e3)) | ~ spl66_9), inference(backward_demodulation, [], [f281, f626])).
fof(f2766, plain, (~ spl66_14 | ~ spl66_16), inference(avatar_contradiction_clause, [], [f2765])).
fof(f2765, plain, ($false | (~ spl66_14 | ~ spl66_16)), inference(subsumption_resolution, [], [f2764, f287])).
fof(f2764, plain, ((e1 = e3) | (~ spl66_14 | ~ spl66_16)), inference(forward_demodulation, [], [f655, f647])).
fof(f2763, plain, (~ spl66_3 | ~ spl66_19), inference(avatar_split_clause, [], [f2761, f666, f598])).
fof(f2761, plain, (~ (e2 = op(e3, e3)) | ~ spl66_19), inference(backward_demodulation, [], [f258, f668])).
fof(f668, plain, ((e2 = op(e2, e3)) | ~ spl66_19), inference(avatar_component_clause, [], [f666])).
fof(f2752, plain, (~ spl66_29 | ~ spl66_32), inference(avatar_contradiction_clause, [], [f2751])).
fof(f2751, plain, ($false | (~ spl66_29 | ~ spl66_32)), inference(subsumption_resolution, [], [f2750, f285])).
fof(f2750, plain, ((e0 = e3) | (~ spl66_29 | ~ spl66_32)), inference(forward_demodulation, [], [f723, f711])).
fof(f2735, plain, (~ spl66_40 | ~ spl66_44), inference(avatar_split_clause, [], [f2728, f772, f755])).
fof(f2728, plain, (~ (e3 = op(e1, e2)) | ~ spl66_44), inference(backward_demodulation, [], [f268, f774])).
fof(f2734, plain, (~ spl66_12 | ~ spl66_44), inference(avatar_split_clause, [], [f2727, f772, f636])).
fof(f2727, plain, (~ (e3 = op(e3, e1)) | ~ spl66_44), inference(backward_demodulation, [], [f245, f774])).
fof(f2733, plain, (~ spl66_28 | ~ spl66_44), inference(avatar_split_clause, [], [f2726, f772, f704])).
fof(f2726, plain, (~ (e3 = op(e2, e1)) | ~ spl66_44), inference(backward_demodulation, [], [f244, f774])).
fof(f2719, plain, (~ spl66_49 | ~ spl66_52), inference(avatar_contradiction_clause, [], [f2718])).
fof(f2718, plain, ($false | (~ spl66_49 | ~ spl66_52)), inference(subsumption_resolution, [], [f2717, f285])).
fof(f2717, plain, ((e0 = e3) | (~ spl66_49 | ~ spl66_52)), inference(forward_demodulation, [], [f808, f796])).
fof(f2685, plain, (~ spl66_46 | ~ spl66_14), inference(avatar_split_clause, [], [f2684, f645, f781])).
fof(f2684, plain, (~ (e1 = op(e1, e0)) | ~ spl66_14), inference(forward_demodulation, [], [f239, f647])).
fof(f2677, plain, (~ spl66_1 | ~ spl66_49), inference(avatar_split_clause, [], [f2676, f794, f590])).
fof(f2676, plain, (~ (e0 = op(e3, e3)) | ~ spl66_49), inference(forward_demodulation, [], [f255, f796])).
fof(f2666, plain, (spl66_2 | ~ spl66_12 | ~ spl66_145), inference(avatar_contradiction_clause, [], [f2665])).
fof(f2665, plain, ($false | (spl66_2 | ~ spl66_12 | ~ spl66_145)), inference(subsumption_resolution, [], [f2663, f595])).
fof(f595, plain, (~ (e1 = op(e3, e3)) | spl66_2), inference(avatar_component_clause, [], [f594])).
fof(f2663, plain, ((e1 = op(e3, e3)) | (~ spl66_12 | ~ spl66_145)), inference(backward_demodulation, [], [f1485, f638])).
fof(f638, plain, ((e3 = op(e3, e1)) | ~ spl66_12), inference(avatar_component_clause, [], [f636])).
fof(f2638, plain, (~ spl66_29 | ~ spl66_31), inference(avatar_contradiction_clause, [], [f2637])).
fof(f2637, plain, ($false | (~ spl66_29 | ~ spl66_31)), inference(subsumption_resolution, [], [f2636, f284])).
fof(f284, plain, ~ (e0 = e2), inference(cnf_transformation, [], [f4])).
fof(f2636, plain, ((e0 = e2) | (~ spl66_29 | ~ spl66_31)), inference(backward_demodulation, [], [f719, f711])).
fof(f2634, plain, (~ spl66_21 | ~ spl66_31 | spl66_132), inference(avatar_split_clause, [], [f2631, f1417, f717, f675])).
fof(f2631, plain, (~ (e0 = op(e2, e2)) | (~ spl66_31 | spl66_132)), inference(backward_demodulation, [], [f1419, f719])).
fof(f2627, plain, (~ spl66_34 | ~ spl66_36), inference(avatar_contradiction_clause, [], [f2626])).
fof(f2626, plain, ($false | (~ spl66_34 | ~ spl66_36)), inference(subsumption_resolution, [], [f2625, f287])).
fof(f2625, plain, ((e1 = e3) | (~ spl66_34 | ~ spl66_36)), inference(forward_demodulation, [], [f740, f732])).
fof(f732, plain, ((e1 = op(e1, e3)) | ~ spl66_34), inference(avatar_component_clause, [], [f730])).
fof(f2622, plain, (~ spl66_41 | ~ spl66_43), inference(avatar_contradiction_clause, [], [f2621])).
fof(f2621, plain, ($false | (~ spl66_41 | ~ spl66_43)), inference(subsumption_resolution, [], [f2620, f284])).
fof(f2620, plain, ((e0 = e2) | (~ spl66_41 | ~ spl66_43)), inference(backward_demodulation, [], [f770, f762])).
fof(f2618, plain, (~ spl66_27 | ~ spl66_43), inference(avatar_split_clause, [], [f2611, f768, f700])).
fof(f2611, plain, (~ (e2 = op(e2, e1)) | ~ spl66_43), inference(backward_demodulation, [], [f244, f770])).
fof(f2610, plain, (~ spl66_38 | ~ spl66_46), inference(avatar_split_clause, [], [f2605, f781, f747])).
fof(f2605, plain, (~ (e1 = op(e1, e2)) | ~ spl66_46), inference(backward_demodulation, [], [f266, f783])).
fof(f2594, plain, (~ spl66_58 | spl66_141), inference(avatar_contradiction_clause, [], [f2593])).
fof(f2593, plain, ($false | (~ spl66_58 | spl66_141)), inference(subsumption_resolution, [], [f2590, f834])).
fof(f2590, plain, (~ (e1 = op(e0, e1)) | (~ spl66_58 | spl66_141)), inference(backward_demodulation, [], [f1462, f834])).
fof(f2592, plain, (~ spl66_50 | ~ spl66_58), inference(avatar_split_clause, [], [f2587, f832, f798])).
fof(f2587, plain, (~ (e1 = op(e0, e3)) | ~ spl66_58), inference(backward_demodulation, [], [f263, f834])).
fof(f2591, plain, (~ spl66_10 | ~ spl66_58), inference(avatar_split_clause, [], [f2586, f832, f628])).
fof(f2586, plain, (~ (e1 = op(e3, e1)) | ~ spl66_58), inference(backward_demodulation, [], [f243, f834])).
fof(f2583, plain, (~ spl66_55 | spl66_140), inference(avatar_contradiction_clause, [], [f2582])).
fof(f2582, plain, ($false | (~ spl66_55 | spl66_140)), inference(subsumption_resolution, [], [f2581, f821])).
fof(f2581, plain, (~ (e2 = op(e0, e2)) | (~ spl66_55 | spl66_140)), inference(forward_demodulation, [], [f1457, f821])).
fof(f1457, plain, (~ (e2 = op(e0, op(e0, e2))) | spl66_140), inference(avatar_component_clause, [], [f1455])).
fof(f1455, plain, (spl66_140 <=> (e2 = op(e0, op(e0, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl66_140])])).
fof(f2574, plain, (~ spl66_38 | ~ spl66_34), inference(avatar_split_clause, [], [f2573, f730, f747])).
fof(f2573, plain, (~ (e1 = op(e1, e2)) | ~ spl66_34), inference(forward_demodulation, [], [f270, f732])).
fof(f2561, plain, (~ spl66_30 | ~ spl66_22), inference(avatar_split_clause, [], [f2560, f679, f713])).
fof(f2560, plain, (~ (e1 = op(e2, e0)) | ~ spl66_22), inference(forward_demodulation, [], [f272, f681])).
fof(f2556, plain, (~ spl66_40 | ~ spl66_22 | spl66_150), inference(avatar_split_clause, [], [f2555, f1511, f679, f755])).
fof(f2555, plain, (~ (e3 = op(e1, e2)) | (~ spl66_22 | spl66_150)), inference(forward_demodulation, [], [f1513, f681])).
fof(f1513, plain, (~ (e3 = op(op(e2, e2), e2)) | spl66_150), inference(avatar_component_clause, [], [f1511])).
fof(f2536, plain, (spl66_8 | ~ spl66_3 | ~ spl66_143), inference(avatar_split_clause, [], [f2460, f1474, f598, f619])).
fof(f2460, plain, ((e3 = op(e3, e2)) | (~ spl66_3 | ~ spl66_143)), inference(backward_demodulation, [], [f1475, f600])).
fof(f2534, plain, (~ spl66_20 | spl66_129), inference(avatar_contradiction_clause, [], [f2533])).
fof(f2533, plain, ($false | (~ spl66_20 | spl66_129)), inference(subsumption_resolution, [], [f2529, f672])).
fof(f672, plain, ((e3 = op(e2, e3)) | ~ spl66_20), inference(avatar_component_clause, [], [f670])).
fof(f2529, plain, (~ (e3 = op(e2, e3)) | (~ spl66_20 | spl66_129)), inference(backward_demodulation, [], [f1404, f672])).
fof(f2528, plain, (~ spl66_21 | ~ spl66_22), inference(avatar_contradiction_clause, [], [f2527])).
fof(f2527, plain, ($false | (~ spl66_21 | ~ spl66_22)), inference(subsumption_resolution, [], [f2526, f283])).
fof(f2526, plain, ((e0 = e1) | (~ spl66_21 | ~ spl66_22)), inference(forward_demodulation, [], [f681, f677])).
fof(f2500, plain, (~ spl66_52 | ~ spl66_64), inference(avatar_split_clause, [], [f2495, f857, f806])).
fof(f2495, plain, (~ (e3 = op(e0, e3)) | ~ spl66_64), inference(backward_demodulation, [], [f261, f859])).
fof(f2499, plain, (~ spl66_32 | ~ spl66_64), inference(avatar_split_clause, [], [f2494, f857, f721])).
fof(f2494, plain, (~ (e3 = op(e2, e0)) | ~ spl66_64), inference(backward_demodulation, [], [f236, f859])).
fof(f2466, plain, (~ spl66_3 | spl66_17 | ~ spl66_158), inference(avatar_contradiction_clause, [], [f2465])).
fof(f2465, plain, ($false | (~ spl66_3 | spl66_17 | ~ spl66_158)), inference(subsumption_resolution, [], [f2461, f659])).
fof(f659, plain, (~ (e0 = op(e2, e3)) | spl66_17), inference(avatar_component_clause, [], [f658])).
fof(f2461, plain, ((e0 = op(e2, e3)) | (~ spl66_3 | ~ spl66_158)), inference(backward_demodulation, [], [f1551, f600])).
fof(f2456, plain, (~ spl66_10 | spl66_145), inference(avatar_contradiction_clause, [], [f2455])).
fof(f2455, plain, ($false | (~ spl66_10 | spl66_145)), inference(subsumption_resolution, [], [f2450, f630])).
fof(f2450, plain, (~ (e1 = op(e3, e1)) | (~ spl66_10 | spl66_145)), inference(backward_demodulation, [], [f1486, f630])).
fof(f2441, plain, (~ spl66_31 | ~ spl66_32), inference(avatar_contradiction_clause, [], [f2440])).
fof(f2440, plain, ($false | (~ spl66_31 | ~ spl66_32)), inference(subsumption_resolution, [], [f2439, f288])).
fof(f2439, plain, ((e2 = e3) | (~ spl66_31 | ~ spl66_32)), inference(forward_demodulation, [], [f723, f719])).
fof(f2374, plain, (~ spl66_21 | ~ spl66_24), inference(avatar_contradiction_clause, [], [f2373])).
fof(f2373, plain, ($false | (~ spl66_21 | ~ spl66_24)), inference(subsumption_resolution, [], [f2372, f285])).
fof(f2372, plain, ((e0 = e3) | (~ spl66_21 | ~ spl66_24)), inference(forward_demodulation, [], [f689, f677])).
fof(f2339, plain, (~ spl66_35 | ~ spl66_48 | ~ spl66_62 | spl66_165), inference(avatar_split_clause, [], [f2335, f1585, f849, f789, f734])).
fof(f2335, plain, (~ (e2 = op(e1, e3)) | (~ spl66_48 | ~ spl66_62 | spl66_165)), inference(backward_demodulation, [], [f2318, f791])).
fof(f2318, plain, (~ (e2 = op(e1, op(e1, e0))) | (~ spl66_62 | spl66_165)), inference(backward_demodulation, [], [f1587, f851])).
fof(f2338, plain, (~ spl66_33 | ~ spl66_48 | spl66_137), inference(avatar_split_clause, [], [f2333, f1441, f789, f726])).
fof(f2333, plain, (~ (e0 = op(e1, e3)) | (~ spl66_48 | spl66_137)), inference(backward_demodulation, [], [f1443, f791])).
fof(f2325, plain, (~ spl66_57 | ~ spl66_60), inference(avatar_contradiction_clause, [], [f2324])).
fof(f2324, plain, ($false | (~ spl66_57 | ~ spl66_60)), inference(subsumption_resolution, [], [f2322, f285])).
fof(f2322, plain, ((e0 = e3) | (~ spl66_57 | ~ spl66_60)), inference(backward_demodulation, [], [f842, f830])).
fof(f2303, plain, (~ spl66_64 | ~ spl66_60), inference(avatar_split_clause, [], [f2302, f840, f857])).
fof(f2302, plain, (~ (op(e0, e0) = e3) | ~ spl66_60), inference(forward_demodulation, [], [f259, f842])).
fof(f2289, plain, (~ spl66_12 | ~ spl66_2 | spl66_143), inference(avatar_split_clause, [], [f2281, f1474, f594, f636])).
fof(f2281, plain, (~ (e3 = op(e3, e1)) | (~ spl66_2 | spl66_143)), inference(backward_demodulation, [], [f1476, f596])).
fof(f2247, plain, (~ spl66_42 | ~ spl66_43), inference(avatar_contradiction_clause, [], [f2246])).
fof(f2246, plain, ($false | (~ spl66_42 | ~ spl66_43)), inference(subsumption_resolution, [], [f2245, f286])).
fof(f286, plain, ~ (e1 = e2), inference(cnf_transformation, [], [f4])).
fof(f2245, plain, ((e1 = e2) | (~ spl66_42 | ~ spl66_43)), inference(backward_demodulation, [], [f770, f766])).
fof(f2243, plain, (~ spl66_13 | ~ spl66_45), inference(avatar_split_clause, [], [f2238, f777, f641])).
fof(f2238, plain, (~ (e0 = op(e3, e0)) | ~ spl66_45), inference(backward_demodulation, [], [f239, f779])).
fof(f2231, plain, (~ spl66_54 | ~ spl66_55), inference(avatar_contradiction_clause, [], [f2230])).
fof(f2230, plain, ($false | (~ spl66_54 | ~ spl66_55)), inference(subsumption_resolution, [], [f2229, f286])).
fof(f2229, plain, ((e1 = e2) | (~ spl66_54 | ~ spl66_55)), inference(backward_demodulation, [], [f821, f817])).
fof(f2218, plain, (~ spl66_63 | ~ spl66_64), inference(avatar_contradiction_clause, [], [f2217])).
fof(f2217, plain, ($false | (~ spl66_63 | ~ spl66_64)), inference(subsumption_resolution, [], [f2216, f288])).
fof(f2216, plain, ((e2 = e3) | (~ spl66_63 | ~ spl66_64)), inference(forward_demodulation, [], [f859, f855])).
fof(f2189, plain, (~ spl66_54 | ~ spl66_21 | spl66_154), inference(avatar_split_clause, [], [f2188, f1530, f675, f815])).
fof(f2188, plain, (~ (e1 = op(e0, e2)) | (~ spl66_21 | spl66_154)), inference(forward_demodulation, [], [f1532, f677])).
fof(f1532, plain, (~ (e1 = op(op(e2, e2), e2)) | spl66_154), inference(avatar_component_clause, [], [f1530])).
fof(f2187, plain, (~ spl66_56 | ~ spl66_21 | spl66_150), inference(avatar_split_clause, [], [f2186, f1511, f675, f823])).
fof(f2186, plain, (~ (e3 = op(e0, e2)) | (~ spl66_21 | spl66_150)), inference(forward_demodulation, [], [f1513, f677])).
fof(f2176, plain, (~ spl66_5 | ~ spl66_13), inference(avatar_split_clause, [], [f2173, f641, f607])).
fof(f2173, plain, (~ (e0 = op(e3, e2)) | ~ spl66_13), inference(backward_demodulation, [], [f278, f643])).
fof(f2175, plain, (~ spl66_13 | spl66_146), inference(avatar_contradiction_clause, [], [f2174])).
fof(f2174, plain, ($false | (~ spl66_13 | spl66_146)), inference(subsumption_resolution, [], [f2171, f643])).
fof(f2171, plain, (~ (e0 = op(e3, e0)) | (~ spl66_13 | spl66_146)), inference(backward_demodulation, [], [f1491, f643])).
fof(f1491, plain, (~ (e0 = op(e3, op(e3, e0))) | spl66_146), inference(avatar_component_clause, [], [f1489])).
fof(f2168, plain, (~ spl66_21 | ~ spl66_23), inference(avatar_contradiction_clause, [], [f2167])).
fof(f2167, plain, ($false | (~ spl66_21 | ~ spl66_23)), inference(subsumption_resolution, [], [f2166, f284])).
fof(f2166, plain, ((e0 = e2) | (~ spl66_21 | ~ spl66_23)), inference(backward_demodulation, [], [f685, f677])).
fof(f2151, plain, (~ spl66_62 | ~ spl66_63), inference(avatar_contradiction_clause, [], [f2150])).
fof(f2150, plain, ($false | (~ spl66_62 | ~ spl66_63)), inference(subsumption_resolution, [], [f2149, f286])).
fof(f2149, plain, ((e1 = e2) | (~ spl66_62 | ~ spl66_63)), inference(forward_demodulation, [], [f855, f851])).
fof(f2125, plain, (~ spl66_62 | ~ spl66_64), inference(avatar_contradiction_clause, [], [f2124])).
fof(f2124, plain, ($false | (~ spl66_62 | ~ spl66_64)), inference(subsumption_resolution, [], [f2123, f287])).
fof(f2123, plain, ((e1 = e3) | (~ spl66_62 | ~ spl66_64)), inference(forward_demodulation, [], [f859, f851])).
fof(f2099, plain, (~ spl66_6 | ~ spl66_38), inference(avatar_split_clause, [], [f2098, f747, f611])).
fof(f2098, plain, (~ (e1 = op(e3, e2)) | ~ spl66_38), inference(forward_demodulation, [], [f251, f749])).
fof(f2093, plain, (~ spl66_30 | ~ spl66_31), inference(avatar_contradiction_clause, [], [f2092])).
fof(f2092, plain, ($false | (~ spl66_30 | ~ spl66_31)), inference(subsumption_resolution, [], [f2091, f286])).
fof(f2091, plain, ((e1 = e2) | (~ spl66_30 | ~ spl66_31)), inference(backward_demodulation, [], [f719, f715])).
fof(f2088, plain, (~ spl66_32 | ~ spl66_25 | ~ spl66_43 | spl66_169), inference(avatar_split_clause, [], [f2087, f1605, f768, f692, f721])).
fof(f2087, plain, (~ (e3 = op(e2, e0)) | (~ spl66_25 | ~ spl66_43 | spl66_169)), inference(forward_demodulation, [], [f2083, f694])).
fof(f2083, plain, (~ (e3 = op(e2, op(e2, e1))) | (~ spl66_43 | spl66_169)), inference(backward_demodulation, [], [f1607, f770])).
fof(f2078, plain, (~ spl66_44 | ~ spl66_48), inference(avatar_split_clause, [], [f2076, f789, f772])).
fof(f2076, plain, (~ (e3 = op(e1, e1)) | ~ spl66_48), inference(backward_demodulation, [], [f265, f791])).
fof(f2073, plain, (~ spl66_55 | ~ spl66_56), inference(avatar_contradiction_clause, [], [f2072])).
fof(f2072, plain, ($false | (~ spl66_55 | ~ spl66_56)), inference(subsumption_resolution, [], [f2071, f288])).
fof(f2071, plain, ((e2 = e3) | (~ spl66_55 | ~ spl66_56)), inference(backward_demodulation, [], [f825, f821])).
fof(f2070, plain, (~ spl66_61 | ~ spl66_62), inference(avatar_contradiction_clause, [], [f2069])).
fof(f2069, plain, ($false | (~ spl66_61 | ~ spl66_62)), inference(subsumption_resolution, [], [f2068, f283])).
fof(f2068, plain, ((e0 = e1) | (~ spl66_61 | ~ spl66_62)), inference(forward_demodulation, [], [f851, f847])).
fof(f2067, plain, (~ spl66_61 | ~ spl66_64), inference(avatar_contradiction_clause, [], [f2066])).
fof(f2066, plain, ($false | (~ spl66_61 | ~ spl66_64)), inference(subsumption_resolution, [], [f2065, f285])).
fof(f2065, plain, ((e0 = e3) | (~ spl66_61 | ~ spl66_64)), inference(forward_demodulation, [], [f859, f847])).
fof(f2034, plain, (~ spl66_18 | ~ spl66_3 | spl66_151), inference(avatar_split_clause, [], [f2023, f1516, f598, f662])).
fof(f2023, plain, (~ (e1 = op(e2, e3)) | (~ spl66_3 | spl66_151)), inference(backward_demodulation, [], [f1518, f600])).
fof(f2029, plain, (~ spl66_3 | ~ spl66_18 | ~ spl66_25 | spl66_147), inference(avatar_contradiction_clause, [], [f2028])).
fof(f2028, plain, ($false | (~ spl66_3 | ~ spl66_18 | ~ spl66_25 | spl66_147)), inference(subsumption_resolution, [], [f2027, f694])).
fof(f2027, plain, (~ (e0 = op(e2, e1)) | (~ spl66_3 | ~ spl66_18 | spl66_147)), inference(forward_demodulation, [], [f2022, f664])).
fof(f2022, plain, (~ (e0 = op(e2, op(e2, e3))) | (~ spl66_3 | spl66_147)), inference(backward_demodulation, [], [f1500, f600])).
fof(f1500, plain, (~ (e0 = op(op(e3, e3), op(op(e3, e3), e3))) | spl66_147), inference(avatar_component_clause, [], [f1498])).
fof(f1498, plain, (spl66_147 <=> (e0 = op(op(e3, e3), op(op(e3, e3), e3)))), introduced(avatar_definition, [new_symbols(naming, [spl66_147])])).
fof(f2010, plain, (~ spl66_2 | ~ spl66_14), inference(avatar_split_clause, [], [f2009, f645, f594])).
fof(f2009, plain, (~ (e1 = op(e3, e3)) | ~ spl66_14), inference(backward_demodulation, [], [f279, f647])).
fof(f1985, plain, (~ spl66_42 | ~ spl66_44), inference(avatar_contradiction_clause, [], [f1984])).
fof(f1984, plain, ($false | (~ spl66_42 | ~ spl66_44)), inference(subsumption_resolution, [], [f1983, f287])).
fof(f1983, plain, ((e1 = e3) | (~ spl66_42 | ~ spl66_44)), inference(backward_demodulation, [], [f774, f766])).
fof(f1961, plain, (~ spl66_42 | ~ spl66_58), inference(avatar_split_clause, [], [f1960, f832, f764])).
fof(f1960, plain, (~ (e1 = op(e1, e1)) | ~ spl66_58), inference(forward_demodulation, [], [f241, f834])).
fof(f1959, plain, (~ spl66_39 | ~ spl66_23), inference(avatar_split_clause, [], [f1958, f683, f751])).
fof(f1958, plain, (~ (e2 = op(e1, e2)) | ~ spl66_23), inference(forward_demodulation, [], [f250, f685])).
fof(f1935, plain, (~ spl66_15 | ~ spl66_5 | spl66_144), inference(avatar_split_clause, [], [f1934, f1479, f607, f649])).
fof(f1934, plain, (~ (e2 = op(e3, e0)) | (~ spl66_5 | spl66_144)), inference(forward_demodulation, [], [f1481, f609])).
fof(f1923, plain, (~ spl66_22 | ~ spl66_23), inference(avatar_contradiction_clause, [], [f1922])).
fof(f1922, plain, ($false | (~ spl66_22 | ~ spl66_23)), inference(subsumption_resolution, [], [f1921, f286])).
fof(f1921, plain, ((e1 = e2) | (~ spl66_22 | ~ spl66_23)), inference(backward_demodulation, [], [f685, f681])).
fof(f1920, plain, (~ spl66_23 | ~ spl66_24), inference(avatar_contradiction_clause, [], [f1919])).
fof(f1919, plain, ($false | (~ spl66_23 | ~ spl66_24)), inference(subsumption_resolution, [], [f1918, f288])).
fof(f1918, plain, ((e2 = e3) | (~ spl66_23 | ~ spl66_24)), inference(backward_demodulation, [], [f689, f685])).
fof(f1903, plain, (~ spl66_24 | ~ spl66_40), inference(avatar_split_clause, [], [f1902, f755, f687])).
fof(f1902, plain, (~ (e3 = op(e2, e2)) | ~ spl66_40), inference(backward_demodulation, [], [f250, f757])).
fof(f757, plain, ((e3 = op(e1, e2)) | ~ spl66_40), inference(avatar_component_clause, [], [f755])).
fof(f1882, plain, (~ spl66_57 | ~ spl66_41), inference(avatar_split_clause, [], [f1881, f760, f828])).
fof(f1881, plain, (~ (e0 = op(e0, e1)) | ~ spl66_41), inference(forward_demodulation, [], [f241, f762])).
fof(f1867, plain, (~ spl66_50 | ~ spl66_2), inference(avatar_split_clause, [], [f1866, f594, f798])).
fof(f1866, plain, (~ (e1 = op(e0, e3)) | ~ spl66_2), inference(forward_demodulation, [], [f255, f596])).
fof(f1853, plain, (~ spl66_34 | ~ spl66_2), inference(avatar_split_clause, [], [f1852, f594, f730])).
fof(f1852, plain, (~ (e1 = op(e1, e3)) | ~ spl66_2), inference(forward_demodulation, [], [f257, f596])).
fof(f1850, plain, (~ spl66_37 | ~ spl66_2 | ~ spl66_35 | spl66_147), inference(avatar_split_clause, [], [f1848, f1498, f734, f594, f743])).
fof(f1848, plain, (~ (e0 = op(e1, e2)) | (~ spl66_2 | ~ spl66_35 | spl66_147)), inference(backward_demodulation, [], [f1811, f736])).
fof(f1811, plain, (~ (e0 = op(e1, op(e1, e3))) | (~ spl66_2 | spl66_147)), inference(forward_demodulation, [], [f1500, f596])).
fof(f1837, plain, (~ spl66_21 | ~ spl66_17), inference(avatar_split_clause, [], [f1836, f658, f675])).
fof(f1836, plain, (~ (e0 = op(e2, e2)) | ~ spl66_17), inference(forward_demodulation, [], [f276, f660])).
fof(f1816, plain, (~ spl66_5 | ~ spl66_8), inference(avatar_contradiction_clause, [], [f1815])).
fof(f1815, plain, ($false | (~ spl66_5 | ~ spl66_8)), inference(subsumption_resolution, [], [f1814, f285])).
fof(f1814, plain, ((e0 = e3) | (~ spl66_5 | ~ spl66_8)), inference(forward_demodulation, [], [f621, f609])).
fof(f621, plain, ((e3 = op(e3, e2)) | ~ spl66_8), inference(avatar_component_clause, [], [f619])).
fof(f1807, plain, (~ spl66_33 | ~ spl66_2 | spl66_158), inference(avatar_split_clause, [], [f1806, f1550, f594, f726])).
fof(f1806, plain, (~ (e0 = op(e1, e3)) | (~ spl66_2 | spl66_158)), inference(forward_demodulation, [], [f1552, f596])).
fof(f1805, plain, (~ spl66_2 | ~ spl66_3), inference(avatar_contradiction_clause, [], [f1804])).
fof(f1804, plain, ($false | (~ spl66_2 | ~ spl66_3)), inference(subsumption_resolution, [], [f1803, f286])).
fof(f1803, plain, ((e1 = e2) | (~ spl66_2 | ~ spl66_3)), inference(backward_demodulation, [], [f600, f596])).
fof(f1802, plain, (~ spl66_3 | ~ spl66_4), inference(avatar_contradiction_clause, [], [f1801])).
fof(f1801, plain, ($false | (~ spl66_3 | ~ spl66_4)), inference(subsumption_resolution, [], [f1800, f288])).
fof(f1800, plain, ((e2 = e3) | (~ spl66_3 | ~ spl66_4)), inference(backward_demodulation, [], [f604, f600])).
fof(f1793, plain, (~ spl66_4 | spl66_143), inference(avatar_contradiction_clause, [], [f1792])).
fof(f1792, plain, ($false | (~ spl66_4 | spl66_143)), inference(subsumption_resolution, [], [f1787, f604])).
fof(f1787, plain, (~ (e3 = op(e3, e3)) | (~ spl66_4 | spl66_143)), inference(backward_demodulation, [], [f1476, f604])).
fof(f1783, plain, (~ spl66_6 | ~ spl66_7), inference(avatar_contradiction_clause, [], [f1782])).
fof(f1782, plain, ($false | (~ spl66_6 | ~ spl66_7)), inference(subsumption_resolution, [], [f1781, f286])).
fof(f1781, plain, ((e1 = e2) | (~ spl66_6 | ~ spl66_7)), inference(backward_demodulation, [], [f617, f613])).
fof(f1780, plain, (~ spl66_7 | ~ spl66_8), inference(avatar_contradiction_clause, [], [f1779])).
fof(f1779, plain, ($false | (~ spl66_7 | ~ spl66_8)), inference(subsumption_resolution, [], [f1778, f288])).
fof(f1778, plain, ((e2 = e3) | (~ spl66_7 | ~ spl66_8)), inference(backward_demodulation, [], [f621, f617])).
fof(f1770, plain, (~ spl66_11 | ~ spl66_12), inference(avatar_contradiction_clause, [], [f1769])).
fof(f1769, plain, ($false | (~ spl66_11 | ~ spl66_12)), inference(subsumption_resolution, [], [f1768, f288])).
fof(f1768, plain, ((e2 = e3) | (~ spl66_11 | ~ spl66_12)), inference(backward_demodulation, [], [f638, f634])).
fof(f1758, plain, (~ spl66_15 | ~ spl66_16), inference(avatar_contradiction_clause, [], [f1757])).
fof(f1757, plain, ($false | (~ spl66_15 | ~ spl66_16)), inference(subsumption_resolution, [], [f1756, f288])).
fof(f1756, plain, ((e2 = e3) | (~ spl66_15 | ~ spl66_16)), inference(backward_demodulation, [], [f655, f651])).
fof(f1755, plain, (~ spl66_1 | ~ spl66_16 | spl66_146), inference(avatar_split_clause, [], [f1751, f1489, f653, f590])).
fof(f1751, plain, (~ (e0 = op(e3, e3)) | (~ spl66_16 | spl66_146)), inference(backward_demodulation, [], [f1491, f655])).
fof(f1744, plain, (~ spl66_18 | ~ spl66_19), inference(avatar_contradiction_clause, [], [f1743])).
fof(f1743, plain, ($false | (~ spl66_18 | ~ spl66_19)), inference(subsumption_resolution, [], [f1742, f286])).
fof(f1742, plain, ((e1 = e2) | (~ spl66_18 | ~ spl66_19)), inference(backward_demodulation, [], [f668, f664])).
fof(f1741, plain, (~ spl66_19 | ~ spl66_20), inference(avatar_contradiction_clause, [], [f1740])).
fof(f1740, plain, ($false | (~ spl66_19 | ~ spl66_20)), inference(subsumption_resolution, [], [f1739, f288])).
fof(f1739, plain, ((e2 = e3) | (~ spl66_19 | ~ spl66_20)), inference(backward_demodulation, [], [f672, f668])).
fof(f1724, plain, (~ spl66_21 | ~ spl66_25), inference(avatar_split_clause, [], [f1721, f692, f675])).
fof(f1721, plain, (~ (e0 = op(e2, e2)) | ~ spl66_25), inference(backward_demodulation, [], [f274, f694])).
fof(f1718, plain, (~ spl66_21 | ~ spl66_29), inference(avatar_split_clause, [], [f1714, f709, f675])).
fof(f1714, plain, (~ (e0 = op(e2, e2)) | ~ spl66_29), inference(backward_demodulation, [], [f272, f711])).
fof(f1716, plain, (~ spl66_13 | ~ spl66_29), inference(avatar_split_clause, [], [f1712, f709, f641])).
fof(f1712, plain, (~ (e0 = op(e3, e0)) | ~ spl66_29), inference(backward_demodulation, [], [f240, f711])).
fof(f1711, plain, (~ spl66_34 | ~ spl66_35), inference(avatar_contradiction_clause, [], [f1710])).
fof(f1710, plain, ($false | (~ spl66_34 | ~ spl66_35)), inference(subsumption_resolution, [], [f1709, f286])).
fof(f1709, plain, ((e1 = e2) | (~ spl66_34 | ~ spl66_35)), inference(backward_demodulation, [], [f736, f732])).
fof(f1708, plain, (~ spl66_35 | ~ spl66_36), inference(avatar_contradiction_clause, [], [f1707])).
fof(f1707, plain, ($false | (~ spl66_35 | ~ spl66_36)), inference(subsumption_resolution, [], [f1706, f288])).
fof(f1706, plain, ((e2 = e3) | (~ spl66_35 | ~ spl66_36)), inference(backward_demodulation, [], [f740, f736])).
fof(f1701, plain, (~ spl66_33 | ~ spl66_37), inference(avatar_split_clause, [], [f1698, f743, f726])).
fof(f1698, plain, (~ (e0 = op(e1, e3)) | ~ spl66_37), inference(backward_demodulation, [], [f270, f745])).
fof(f745, plain, ((e0 = op(e1, e2)) | ~ spl66_37), inference(avatar_component_clause, [], [f743])).
fof(f1699, plain, (~ spl66_21 | ~ spl66_37), inference(avatar_split_clause, [], [f1696, f743, f675])).
fof(f1696, plain, (~ (e0 = op(e2, e2)) | ~ spl66_37), inference(backward_demodulation, [], [f250, f745])).
fof(f1687, plain, (~ spl66_37 | ~ spl66_41), inference(avatar_split_clause, [], [f1680, f760, f743])).
fof(f1680, plain, (~ (e0 = op(e1, e2)) | ~ spl66_41), inference(backward_demodulation, [], [f268, f762])).
fof(f1686, plain, (~ spl66_9 | ~ spl66_41), inference(avatar_split_clause, [], [f1679, f760, f624])).
fof(f1679, plain, (~ (e0 = op(e3, e1)) | ~ spl66_41), inference(backward_demodulation, [], [f245, f762])).
fof(f1667, plain, (~ spl66_50 | ~ spl66_51), inference(avatar_contradiction_clause, [], [f1666])).
fof(f1666, plain, ($false | (~ spl66_50 | ~ spl66_51)), inference(subsumption_resolution, [], [f1665, f286])).
fof(f1665, plain, ((e1 = e2) | (~ spl66_50 | ~ spl66_51)), inference(backward_demodulation, [], [f804, f800])).
fof(f1664, plain, (~ spl66_51 | ~ spl66_52), inference(avatar_contradiction_clause, [], [f1663])).
fof(f1663, plain, ($false | (~ spl66_51 | ~ spl66_52)), inference(subsumption_resolution, [], [f1662, f288])).
fof(f1662, plain, ((e2 = e3) | (~ spl66_51 | ~ spl66_52)), inference(backward_demodulation, [], [f808, f804])).
fof(f1655, plain, (~ spl66_49 | ~ spl66_53), inference(avatar_split_clause, [], [f1651, f811, f794])).
fof(f1651, plain, (~ (e0 = op(e0, e3)) | ~ spl66_53), inference(backward_demodulation, [], [f264, f813])).
fof(f1630, plain, (~ spl66_53 | ~ spl66_61), inference(avatar_split_clause, [], [f1621, f845, f811])).
fof(f1621, plain, (~ (e0 = op(e0, e2)) | ~ spl66_61), inference(backward_demodulation, [], [f260, f847])).
fof(f1626, plain, (~ spl66_45 | ~ spl66_61), inference(avatar_split_clause, [], [f1617, f845, f777])).
fof(f1617, plain, (~ (e0 = op(e1, e0)) | ~ spl66_61), inference(backward_demodulation, [], [f235, f847])).
fof(f1616, plain, (~ spl66_170 | ~ spl66_63 | ~ spl66_167), inference(avatar_split_clause, [], [f588, f1595, f853, f1611])).
fof(f588, plain, (~ (e1 = op(op(e0, e0), e0)) | ~ (op(e0, e0) = e2) | ~ (e3 = op(op(e0, e0), op(op(e0, e0), e0)))), inference(cnf_transformation, [], [f53])).
fof(f53, plain, (~ (e1 = op(op(e0, e0), e0)) | ~ (op(e0, e0) = e2) | ~ (e3 = op(op(e0, e0), op(op(e0, e0), e0)))), inference(ennf_transformation, [], [f29])).
fof(f29, plain, ~ ((e1 = op(op(e0, e0), e0)) & (op(e0, e0) = e2) & (e3 = op(op(e0, e0), op(op(e0, e0), e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG140+1.p', ax29)).
fof(f1615, plain, (~ spl66_169 | ~ spl66_43 | ~ spl66_166), inference(avatar_split_clause, [], [f587, f1590, f768, f1605])).
fof(f587, plain, (~ (e0 = op(op(e1, e1), e1)) | ~ (e2 = op(e1, e1)) | ~ (e3 = op(op(e1, e1), op(op(e1, e1), e1)))), inference(cnf_transformation, [], [f52])).
fof(f52, plain, (~ (e0 = op(op(e1, e1), e1)) | ~ (e2 = op(e1, e1)) | ~ (e3 = op(op(e1, e1), op(op(e1, e1), e1)))), inference(ennf_transformation, [], [f28])).
fof(f28, plain, ~ ((e0 = op(op(e1, e1), e1)) & (e2 = op(e1, e1)) & (e3 = op(op(e1, e1), op(op(e1, e1), e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG140+1.p', ax28)).
fof(f1614, plain, (~ spl66_170 | ~ spl66_62 | ~ spl66_162), inference(avatar_split_clause, [], [f586, f1569, f849, f1611])).
fof(f586, plain, (~ (e2 = op(op(e0, e0), e0)) | ~ (op(e0, e0) = e1) | ~ (e3 = op(op(e0, e0), op(op(e0, e0), e0)))), inference(cnf_transformation, [], [f51])).
fof(f51, plain, (~ (e2 = op(op(e0, e0), e0)) | ~ (op(e0, e0) = e1) | ~ (e3 = op(op(e0, e0), op(op(e0, e0), e0)))), inference(ennf_transformation, [], [f27])).
fof(f27, plain, ~ ((e2 = op(op(e0, e0), e0)) & (op(e0, e0) = e1) & (e3 = op(op(e0, e0), op(op(e0, e0), e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG140+1.p', ax27)).
fof(f1609, plain, (~ spl66_168 | ~ spl66_22 | ~ spl66_161), inference(avatar_split_clause, [], [f585, f1564, f679, f1600])).
fof(f585, plain, (~ (e0 = op(op(e2, e2), e2)) | ~ (e1 = op(e2, e2)) | ~ (e3 = op(op(e2, e2), op(op(e2, e2), e2)))), inference(cnf_transformation, [], [f50])).
fof(f50, plain, (~ (e0 = op(op(e2, e2), e2)) | ~ (e1 = op(e2, e2)) | ~ (e3 = op(op(e2, e2), op(op(e2, e2), e2)))), inference(ennf_transformation, [], [f26])).
fof(f26, plain, ~ ((e0 = op(op(e2, e2), e2)) & (e1 = op(e2, e2)) & (e3 = op(op(e2, e2), op(op(e2, e2), e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG140+1.p', ax26)).
fof(f1603, plain, (~ spl66_168 | ~ spl66_21 | ~ spl66_154), inference(avatar_split_clause, [], [f583, f1530, f675, f1600])).
fof(f583, plain, (~ (e1 = op(op(e2, e2), e2)) | ~ (e0 = op(e2, e2)) | ~ (e3 = op(op(e2, e2), op(op(e2, e2), e2)))), inference(cnf_transformation, [], [f48])).
fof(f48, plain, (~ (e1 = op(op(e2, e2), e2)) | ~ (e0 = op(e2, e2)) | ~ (e3 = op(op(e2, e2), op(op(e2, e2), e2)))), inference(ennf_transformation, [], [f24])).
fof(f24, plain, ~ ((e1 = op(op(e2, e2), e2)) & (e0 = op(e2, e2)) & (e3 = op(op(e2, e2), op(op(e2, e2), e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG140+1.p', ax24)).
fof(f1598, plain, (~ spl66_165 | ~ spl66_64 | ~ spl66_167), inference(avatar_split_clause, [], [f582, f1595, f857, f1585])).
fof(f582, plain, (~ (e1 = op(op(e0, e0), e0)) | ~ (op(e0, e0) = e3) | ~ (e2 = op(op(e0, e0), op(op(e0, e0), e0)))), inference(cnf_transformation, [], [f47])).
fof(f47, plain, (~ (e1 = op(op(e0, e0), e0)) | ~ (op(e0, e0) = e3) | ~ (e2 = op(op(e0, e0), op(op(e0, e0), e0)))), inference(ennf_transformation, [], [f23])).
fof(f23, plain, ~ ((e1 = op(op(e0, e0), e0)) & (op(e0, e0) = e3) & (e2 = op(op(e0, e0), op(op(e0, e0), e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG140+1.p', ax23)).
fof(f1588, plain, (~ spl66_165 | ~ spl66_62 | ~ spl66_160), inference(avatar_split_clause, [], [f580, f1559, f849, f1585])).
fof(f580, plain, (~ (e3 = op(op(e0, e0), e0)) | ~ (op(e0, e0) = e1) | ~ (e2 = op(op(e0, e0), op(op(e0, e0), e0)))), inference(cnf_transformation, [], [f45])).
fof(f45, plain, (~ (e3 = op(op(e0, e0), e0)) | ~ (op(e0, e0) = e1) | ~ (e2 = op(op(e0, e0), op(op(e0, e0), e0)))), inference(ennf_transformation, [], [f21])).
fof(f21, plain, ~ ((e3 = op(op(e0, e0), e0)) & (op(e0, e0) = e1) & (e2 = op(op(e0, e0), op(op(e0, e0), e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG140+1.p', ax21)).
fof(f1583, plain, (~ spl66_163 | ~ spl66_2 | ~ spl66_158), inference(avatar_split_clause, [], [f579, f1550, f594, f1574])).
fof(f579, plain, (~ (e0 = op(op(e3, e3), e3)) | ~ (e1 = op(e3, e3)) | ~ (e2 = op(op(e3, e3), op(op(e3, e3), e3)))), inference(cnf_transformation, [], [f44])).
fof(f44, plain, (~ (e0 = op(op(e3, e3), e3)) | ~ (e1 = op(e3, e3)) | ~ (e2 = op(op(e3, e3), op(op(e3, e3), e3)))), inference(ennf_transformation, [], [f20])).
fof(f20, plain, ~ ((e0 = op(op(e3, e3), e3)) & (e1 = op(e3, e3)) & (e2 = op(op(e3, e3), op(op(e3, e3), e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG140+1.p', ax20)).
fof(f1572, plain, (~ spl66_159 | ~ spl66_64 | ~ spl66_162), inference(avatar_split_clause, [], [f576, f1569, f857, f1555])).
fof(f576, plain, (~ (e2 = op(op(e0, e0), e0)) | ~ (op(e0, e0) = e3) | ~ (e1 = op(op(e0, e0), op(op(e0, e0), e0)))), inference(cnf_transformation, [], [f41])).
fof(f41, plain, (~ (e2 = op(op(e0, e0), e0)) | ~ (op(e0, e0) = e3) | ~ (e1 = op(op(e0, e0), op(op(e0, e0), e0)))), inference(ennf_transformation, [], [f17])).
fof(f17, plain, ~ ((e2 = op(op(e0, e0), e0)) & (op(e0, e0) = e3) & (e1 = op(op(e0, e0), op(op(e0, e0), e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG140+1.p', ax17)).
fof(f1553, plain, (~ spl66_156 | ~ spl66_3 | ~ spl66_158), inference(avatar_split_clause, [], [f573, f1550, f598, f1540])).
fof(f573, plain, (~ (e0 = op(op(e3, e3), e3)) | ~ (e2 = op(e3, e3)) | ~ (e1 = op(op(e3, e3), op(op(e3, e3), e3)))), inference(cnf_transformation, [], [f38])).
fof(f38, plain, (~ (e0 = op(op(e3, e3), e3)) | ~ (e2 = op(e3, e3)) | ~ (e1 = op(op(e3, e3), op(op(e3, e3), e3)))), inference(ennf_transformation, [], [f14])).
fof(f14, plain, ~ ((e0 = op(op(e3, e3), e3)) & (e2 = op(e3, e3)) & (e1 = op(op(e3, e3), op(op(e3, e3), e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG140+1.p', ax14)).
fof(f1548, plain, (~ spl66_157 | ~ spl66_21 | ~ spl66_150), inference(avatar_split_clause, [], [f572, f1511, f675, f1545])).
fof(f572, plain, (~ (e3 = op(op(e2, e2), e2)) | ~ (e0 = op(e2, e2)) | ~ (e1 = op(op(e2, e2), op(op(e2, e2), e2)))), inference(cnf_transformation, [], [f37])).
fof(f37, plain, (~ (e3 = op(op(e2, e2), e2)) | ~ (e0 = op(e2, e2)) | ~ (e1 = op(op(e2, e2), op(op(e2, e2), e2)))), inference(ennf_transformation, [], [f13])).
fof(f13, plain, ~ ((e3 = op(op(e2, e2), e2)) & (e0 = op(e2, e2)) & (e1 = op(op(e2, e2), op(op(e2, e2), e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG140+1.p', ax13)).
fof(f1538, plain, (~ spl66_152 | ~ spl66_44 | ~ spl66_155), inference(avatar_split_clause, [], [f570, f1535, f772, f1521])).
fof(f570, plain, (~ (e2 = op(op(e1, e1), e1)) | ~ (e3 = op(e1, e1)) | ~ (e0 = op(op(e1, e1), op(op(e1, e1), e1)))), inference(cnf_transformation, [], [f35])).
fof(f35, plain, (~ (e2 = op(op(e1, e1), e1)) | ~ (e3 = op(e1, e1)) | ~ (e0 = op(op(e1, e1), op(op(e1, e1), e1)))), inference(ennf_transformation, [], [f11])).
fof(f11, plain, ~ ((e2 = op(op(e1, e1), e1)) & (e3 = op(e1, e1)) & (e0 = op(op(e1, e1), op(op(e1, e1), e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG140+1.p', ax11)).
fof(f1519, plain, (~ spl66_147 | ~ spl66_3 | ~ spl66_151), inference(avatar_split_clause, [], [f567, f1516, f598, f1498])).
fof(f567, plain, (~ (e1 = op(op(e3, e3), e3)) | ~ (e2 = op(e3, e3)) | ~ (e0 = op(op(e3, e3), op(op(e3, e3), e3)))), inference(cnf_transformation, [], [f32])).
fof(f32, plain, (~ (e1 = op(op(e3, e3), e3)) | ~ (e2 = op(e3, e3)) | ~ (e0 = op(op(e3, e3), op(op(e3, e3), e3)))), inference(ennf_transformation, [], [f8])).
fof(f8, plain, ~ ((e1 = op(op(e3, e3), e3)) & (e2 = op(e3, e3)) & (e0 = op(op(e3, e3), op(op(e3, e3), e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG140+1.p', ax8)).
fof(f1505, plain, (~ spl66_147 | ~ spl66_2 | ~ spl66_148), inference(avatar_split_clause, [], [f565, f1502, f594, f1498])).
fof(f565, plain, (~ (e2 = op(op(e3, e3), e3)) | ~ (e1 = op(e3, e3)) | ~ (e0 = op(op(e3, e3), op(op(e3, e3), e3)))), inference(cnf_transformation, [], [f30])).
fof(f30, plain, (~ (e2 = op(op(e3, e3), e3)) | ~ (e1 = op(e3, e3)) | ~ (e0 = op(op(e3, e3), op(op(e3, e3), e3)))), inference(ennf_transformation, [], [f6])).
fof(f6, plain, ~ ((e2 = op(op(e3, e3), e3)) & (e1 = op(e3, e3)) & (e0 = op(op(e3, e3), op(op(e3, e3), e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG140+1.p', ax6)).
fof(f1496, plain, (spl66_61 | spl66_41 | spl66_21 | spl66_1), inference(avatar_split_clause, [], [f553, f590, f675, f760, f845])).
fof(f553, plain, ((e0 = op(e3, e3)) | (e0 = op(e2, e2)) | (e0 = op(e1, e1)) | (e0 = op(e0, e0))), inference(cnf_transformation, [], [f120])).
fof(f120, plain, ((((e3 = op(e3, e3)) & (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3) & ((~ (e3 = op(e3, op(e3, e3))) & ~ (e2 = op(e3, op(e3, e2))) & ~ (e1 = op(e3, op(e3, e1))) & ~ (e0 = op(e3, op(e3, e0)))) | sP2 | sP1 | sP0) & ((e3 = op(e3, e3)) | (e3 = op(e2, e2)) | (e3 = op(e1, e1)) | (op(e0, e0) = e3)) & ((e2 = op(e3, e3)) | (e2 = op(e2, e2)) | (e2 = op(e1, e1)) | (op(e0, e0) = e2)) & ((e1 = op(e3, e3)) | (e1 = op(e2, e2)) | (e1 = op(e1, e1)) | (op(e0, e0) = e1)) & ((e0 = op(e3, e3)) | (e0 = op(e2, e2)) | (e0 = op(e1, e1)) | (e0 = op(e0, e0)))), inference(definition_folding, [], [f5, e119, e118, e117, e116, e115, e114, e113, e112, e111, e110, e109, e108, e107, e106, e105, e104, e103, e102, e101, e100, e99, e98, e97, e96, e95, e94, e93, e92, e91, e90, e89, e88, e87, e86, e85, e84, e83, e82, e81, e80, e79, e78, e77, e76, e75, e74, e73, e72, e71, e70, e69, e68, e67, e66, e65, e64, e63, e62, e61, e60, e59, e58, e57, e56, e55, e54])).
fof(f54, plain, ((~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0)))) | ~ sP0), inference(usedef, [], [e54])).
fof(e54, plain, (sP0 <=> (~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f55, plain, ((~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0)))) | ~ sP1), inference(usedef, [], [e55])).
fof(e55, plain, (sP1 <=> (~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f56, plain, ((~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0)))) | ~ sP2), inference(usedef, [], [e56])).
fof(e56, plain, (sP2 <=> (~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f57, plain, (((e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP3), inference(usedef, [], [e57])).
fof(e57, plain, (sP3 <=> ((e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f58, plain, (((e0 = op(e0, e1)) & (e0 = op(e1, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP4), inference(usedef, [], [e58])).
fof(e58, plain, (sP4 <=> ((e0 = op(e0, e1)) & (e0 = op(e1, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f59, plain, (((e0 = op(e0, e2)) & (e0 = op(e2, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP5), inference(usedef, [], [e59])).
fof(e59, plain, (sP5 <=> ((e0 = op(e0, e2)) & (e0 = op(e2, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f60, plain, (((e0 = op(e0, e3)) & (e0 = op(e3, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP6), inference(usedef, [], [e60])).
fof(e60, plain, (sP6 <=> ((e0 = op(e0, e3)) & (e0 = op(e3, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f61, plain, (((e1 = op(e1, e0)) & (e1 = op(e0, e1)) & (e0 = op(e1, e0)) & (e0 = op(e0, e1))) | ~ sP7), inference(usedef, [], [e61])).
fof(e61, plain, (sP7 <=> ((e1 = op(e1, e0)) & (e1 = op(e0, e1)) & (e0 = op(e1, e0)) & (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f62, plain, (((e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e0 = op(e1, e0)) & (e0 = op(e0, e1))) | ~ sP8), inference(usedef, [], [e62])).
fof(e62, plain, (sP8 <=> ((e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e0 = op(e1, e0)) & (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f63, plain, (((e1 = op(e1, e2)) & (e1 = op(e2, e1)) & (e0 = op(e1, e0)) & (e0 = op(e0, e1))) | ~ sP9), inference(usedef, [], [e63])).
fof(e63, plain, (sP9 <=> ((e1 = op(e1, e2)) & (e1 = op(e2, e1)) & (e0 = op(e1, e0)) & (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f64, plain, (((e1 = op(e1, e3)) & (e1 = op(e3, e1)) & (e0 = op(e1, e0)) & (e0 = op(e0, e1))) | ~ sP10), inference(usedef, [], [e64])).
fof(e64, plain, (sP10 <=> ((e1 = op(e1, e3)) & (e1 = op(e3, e1)) & (e0 = op(e1, e0)) & (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f65, plain, (((e2 = op(e2, e0)) & (e2 = op(e0, e2)) & (e0 = op(e2, e0)) & (e0 = op(e0, e2))) | ~ sP11), inference(usedef, [], [e65])).
fof(e65, plain, (sP11 <=> ((e2 = op(e2, e0)) & (e2 = op(e0, e2)) & (e0 = op(e2, e0)) & (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f66, plain, (((e2 = op(e2, e1)) & (e2 = op(e1, e2)) & (e0 = op(e2, e0)) & (e0 = op(e0, e2))) | ~ sP12), inference(usedef, [], [e66])).
fof(e66, plain, (sP12 <=> ((e2 = op(e2, e1)) & (e2 = op(e1, e2)) & (e0 = op(e2, e0)) & (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f67, plain, (((e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e0 = op(e2, e0)) & (e0 = op(e0, e2))) | ~ sP13), inference(usedef, [], [e67])).
fof(e67, plain, (sP13 <=> ((e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e0 = op(e2, e0)) & (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f68, plain, (((e2 = op(e2, e3)) & (e2 = op(e3, e2)) & (e0 = op(e2, e0)) & (e0 = op(e0, e2))) | ~ sP14), inference(usedef, [], [e68])).
fof(e68, plain, (sP14 <=> ((e2 = op(e2, e3)) & (e2 = op(e3, e2)) & (e0 = op(e2, e0)) & (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f69, plain, (((e3 = op(e3, e0)) & (e3 = op(e0, e3)) & (e0 = op(e3, e0)) & (e0 = op(e0, e3))) | ~ sP15), inference(usedef, [], [e69])).
fof(e69, plain, (sP15 <=> ((e3 = op(e3, e0)) & (e3 = op(e0, e3)) & (e0 = op(e3, e0)) & (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f70, plain, (((e3 = op(e3, e1)) & (e3 = op(e1, e3)) & (e0 = op(e3, e0)) & (e0 = op(e0, e3))) | ~ sP16), inference(usedef, [], [e70])).
fof(e70, plain, (sP16 <=> ((e3 = op(e3, e1)) & (e3 = op(e1, e3)) & (e0 = op(e3, e0)) & (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f71, plain, (((e3 = op(e3, e2)) & (e3 = op(e2, e3)) & (e0 = op(e3, e0)) & (e0 = op(e0, e3))) | ~ sP17), inference(usedef, [], [e71])).
fof(e71, plain, (sP17 <=> ((e3 = op(e3, e2)) & (e3 = op(e2, e3)) & (e0 = op(e3, e0)) & (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f72, plain, (((e3 = op(e3, e3)) & (e3 = op(e3, e3)) & (e0 = op(e3, e0)) & (e0 = op(e0, e3))) | ~ sP18), inference(usedef, [], [e72])).
fof(e72, plain, (sP18 <=> ((e3 = op(e3, e3)) & (e3 = op(e3, e3)) & (e0 = op(e3, e0)) & (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f73, plain, (((e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e1 = op(e0, e1)) & (e1 = op(e1, e0))) | ~ sP19), inference(usedef, [], [e73])).
fof(e73, plain, (sP19 <=> ((e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e1 = op(e0, e1)) & (e1 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP19])])).
fof(f74, plain, (((e0 = op(e0, e1)) & (e0 = op(e1, e0)) & (e1 = op(e0, e1)) & (e1 = op(e1, e0))) | ~ sP20), inference(usedef, [], [e74])).
fof(e74, plain, (sP20 <=> ((e0 = op(e0, e1)) & (e0 = op(e1, e0)) & (e1 = op(e0, e1)) & (e1 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP20])])).
fof(f75, plain, (((e0 = op(e0, e2)) & (e0 = op(e2, e0)) & (e1 = op(e0, e1)) & (e1 = op(e1, e0))) | ~ sP21), inference(usedef, [], [e75])).
fof(e75, plain, (sP21 <=> ((e0 = op(e0, e2)) & (e0 = op(e2, e0)) & (e1 = op(e0, e1)) & (e1 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP21])])).
fof(f76, plain, (((e0 = op(e0, e3)) & (e0 = op(e3, e0)) & (e1 = op(e0, e1)) & (e1 = op(e1, e0))) | ~ sP22), inference(usedef, [], [e76])).
fof(e76, plain, (sP22 <=> ((e0 = op(e0, e3)) & (e0 = op(e3, e0)) & (e1 = op(e0, e1)) & (e1 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP22])])).
fof(f77, plain, (((e1 = op(e1, e0)) & (e1 = op(e0, e1)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP23), inference(usedef, [], [e77])).
fof(e77, plain, (sP23 <=> ((e1 = op(e1, e0)) & (e1 = op(e0, e1)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP23])])).
fof(f78, plain, (((e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP24), inference(usedef, [], [e78])).
fof(e78, plain, (sP24 <=> ((e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP24])])).
fof(f79, plain, (((e1 = op(e1, e2)) & (e1 = op(e2, e1)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP25), inference(usedef, [], [e79])).
fof(e79, plain, (sP25 <=> ((e1 = op(e1, e2)) & (e1 = op(e2, e1)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP25])])).
fof(f80, plain, (((e1 = op(e1, e3)) & (e1 = op(e3, e1)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP26), inference(usedef, [], [e80])).
fof(e80, plain, (sP26 <=> ((e1 = op(e1, e3)) & (e1 = op(e3, e1)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP26])])).
fof(f81, plain, (((e2 = op(e2, e0)) & (e2 = op(e0, e2)) & (e1 = op(e2, e1)) & (e1 = op(e1, e2))) | ~ sP27), inference(usedef, [], [e81])).
fof(e81, plain, (sP27 <=> ((e2 = op(e2, e0)) & (e2 = op(e0, e2)) & (e1 = op(e2, e1)) & (e1 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP27])])).
fof(f82, plain, (((e2 = op(e2, e1)) & (e2 = op(e1, e2)) & (e1 = op(e2, e1)) & (e1 = op(e1, e2))) | ~ sP28), inference(usedef, [], [e82])).
fof(e82, plain, (sP28 <=> ((e2 = op(e2, e1)) & (e2 = op(e1, e2)) & (e1 = op(e2, e1)) & (e1 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP28])])).
fof(f83, plain, (((e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e1 = op(e2, e1)) & (e1 = op(e1, e2))) | ~ sP29), inference(usedef, [], [e83])).
fof(e83, plain, (sP29 <=> ((e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e1 = op(e2, e1)) & (e1 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP29])])).
fof(f84, plain, (((e2 = op(e2, e3)) & (e2 = op(e3, e2)) & (e1 = op(e2, e1)) & (e1 = op(e1, e2))) | ~ sP30), inference(usedef, [], [e84])).
fof(e84, plain, (sP30 <=> ((e2 = op(e2, e3)) & (e2 = op(e3, e2)) & (e1 = op(e2, e1)) & (e1 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP30])])).
fof(f85, plain, (((e3 = op(e3, e0)) & (e3 = op(e0, e3)) & (e1 = op(e3, e1)) & (e1 = op(e1, e3))) | ~ sP31), inference(usedef, [], [e85])).
fof(e85, plain, (sP31 <=> ((e3 = op(e3, e0)) & (e3 = op(e0, e3)) & (e1 = op(e3, e1)) & (e1 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP31])])).
fof(f86, plain, (((e3 = op(e3, e1)) & (e3 = op(e1, e3)) & (e1 = op(e3, e1)) & (e1 = op(e1, e3))) | ~ sP32), inference(usedef, [], [e86])).
fof(e86, plain, (sP32 <=> ((e3 = op(e3, e1)) & (e3 = op(e1, e3)) & (e1 = op(e3, e1)) & (e1 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP32])])).
fof(f87, plain, (((e3 = op(e3, e2)) & (e3 = op(e2, e3)) & (e1 = op(e3, e1)) & (e1 = op(e1, e3))) | ~ sP33), inference(usedef, [], [e87])).
fof(e87, plain, (sP33 <=> ((e3 = op(e3, e2)) & (e3 = op(e2, e3)) & (e1 = op(e3, e1)) & (e1 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP33])])).
fof(f88, plain, (((e3 = op(e3, e3)) & (e3 = op(e3, e3)) & (e1 = op(e3, e1)) & (e1 = op(e1, e3))) | ~ sP34), inference(usedef, [], [e88])).
fof(e88, plain, (sP34 <=> ((e3 = op(e3, e3)) & (e3 = op(e3, e3)) & (e1 = op(e3, e1)) & (e1 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP34])])).
fof(f89, plain, (((e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e2 = op(e0, e2)) & (e2 = op(e2, e0))) | ~ sP35), inference(usedef, [], [e89])).
fof(e89, plain, (sP35 <=> ((e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e2 = op(e0, e2)) & (e2 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP35])])).
fof(f90, plain, (((e0 = op(e0, e1)) & (e0 = op(e1, e0)) & (e2 = op(e0, e2)) & (e2 = op(e2, e0))) | ~ sP36), inference(usedef, [], [e90])).
fof(e90, plain, (sP36 <=> ((e0 = op(e0, e1)) & (e0 = op(e1, e0)) & (e2 = op(e0, e2)) & (e2 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP36])])).
fof(f91, plain, (((e0 = op(e0, e2)) & (e0 = op(e2, e0)) & (e2 = op(e0, e2)) & (e2 = op(e2, e0))) | ~ sP37), inference(usedef, [], [e91])).
fof(e91, plain, (sP37 <=> ((e0 = op(e0, e2)) & (e0 = op(e2, e0)) & (e2 = op(e0, e2)) & (e2 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP37])])).
fof(f92, plain, (((e0 = op(e0, e3)) & (e0 = op(e3, e0)) & (e2 = op(e0, e2)) & (e2 = op(e2, e0))) | ~ sP38), inference(usedef, [], [e92])).
fof(e92, plain, (sP38 <=> ((e0 = op(e0, e3)) & (e0 = op(e3, e0)) & (e2 = op(e0, e2)) & (e2 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP38])])).
fof(f93, plain, (((e1 = op(e1, e0)) & (e1 = op(e0, e1)) & (e2 = op(e1, e2)) & (e2 = op(e2, e1))) | ~ sP39), inference(usedef, [], [e93])).
fof(e93, plain, (sP39 <=> ((e1 = op(e1, e0)) & (e1 = op(e0, e1)) & (e2 = op(e1, e2)) & (e2 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP39])])).
fof(f94, plain, (((e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e2 = op(e1, e2)) & (e2 = op(e2, e1))) | ~ sP40), inference(usedef, [], [e94])).
fof(e94, plain, (sP40 <=> ((e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e2 = op(e1, e2)) & (e2 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP40])])).
fof(f95, plain, (((e1 = op(e1, e2)) & (e1 = op(e2, e1)) & (e2 = op(e1, e2)) & (e2 = op(e2, e1))) | ~ sP41), inference(usedef, [], [e95])).
fof(e95, plain, (sP41 <=> ((e1 = op(e1, e2)) & (e1 = op(e2, e1)) & (e2 = op(e1, e2)) & (e2 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP41])])).
fof(f96, plain, (((e1 = op(e1, e3)) & (e1 = op(e3, e1)) & (e2 = op(e1, e2)) & (e2 = op(e2, e1))) | ~ sP42), inference(usedef, [], [e96])).
fof(e96, plain, (sP42 <=> ((e1 = op(e1, e3)) & (e1 = op(e3, e1)) & (e2 = op(e1, e2)) & (e2 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP42])])).
fof(f97, plain, (((e2 = op(e2, e0)) & (e2 = op(e0, e2)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP43), inference(usedef, [], [e97])).
fof(e97, plain, (sP43 <=> ((e2 = op(e2, e0)) & (e2 = op(e0, e2)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP43])])).
fof(f98, plain, (((e2 = op(e2, e1)) & (e2 = op(e1, e2)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP44), inference(usedef, [], [e98])).
fof(e98, plain, (sP44 <=> ((e2 = op(e2, e1)) & (e2 = op(e1, e2)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP44])])).
fof(f99, plain, (((e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP45), inference(usedef, [], [e99])).
fof(e99, plain, (sP45 <=> ((e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP45])])).
fof(f100, plain, (((e2 = op(e2, e3)) & (e2 = op(e3, e2)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP46), inference(usedef, [], [e100])).
fof(e100, plain, (sP46 <=> ((e2 = op(e2, e3)) & (e2 = op(e3, e2)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP46])])).
fof(f101, plain, (((e3 = op(e3, e0)) & (e3 = op(e0, e3)) & (e2 = op(e3, e2)) & (e2 = op(e2, e3))) | ~ sP47), inference(usedef, [], [e101])).
fof(e101, plain, (sP47 <=> ((e3 = op(e3, e0)) & (e3 = op(e0, e3)) & (e2 = op(e3, e2)) & (e2 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP47])])).
fof(f102, plain, (((e3 = op(e3, e1)) & (e3 = op(e1, e3)) & (e2 = op(e3, e2)) & (e2 = op(e2, e3))) | ~ sP48), inference(usedef, [], [e102])).
fof(e102, plain, (sP48 <=> ((e3 = op(e3, e1)) & (e3 = op(e1, e3)) & (e2 = op(e3, e2)) & (e2 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP48])])).
fof(f103, plain, (((e3 = op(e3, e2)) & (e3 = op(e2, e3)) & (e2 = op(e3, e2)) & (e2 = op(e2, e3))) | ~ sP49), inference(usedef, [], [e103])).
fof(e103, plain, (sP49 <=> ((e3 = op(e3, e2)) & (e3 = op(e2, e3)) & (e2 = op(e3, e2)) & (e2 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP49])])).
fof(f104, plain, (((e3 = op(e3, e3)) & (e3 = op(e3, e3)) & (e2 = op(e3, e2)) & (e2 = op(e2, e3))) | ~ sP50), inference(usedef, [], [e104])).
fof(e104, plain, (sP50 <=> ((e3 = op(e3, e3)) & (e3 = op(e3, e3)) & (e2 = op(e3, e2)) & (e2 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP50])])).
fof(f105, plain, (((e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e3 = op(e0, e3)) & (e3 = op(e3, e0))) | ~ sP51), inference(usedef, [], [e105])).
fof(e105, plain, (sP51 <=> ((e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e3 = op(e0, e3)) & (e3 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP51])])).
fof(f106, plain, (((e0 = op(e0, e1)) & (e0 = op(e1, e0)) & (e3 = op(e0, e3)) & (e3 = op(e3, e0))) | ~ sP52), inference(usedef, [], [e106])).
fof(e106, plain, (sP52 <=> ((e0 = op(e0, e1)) & (e0 = op(e1, e0)) & (e3 = op(e0, e3)) & (e3 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP52])])).
fof(f107, plain, (((e0 = op(e0, e2)) & (e0 = op(e2, e0)) & (e3 = op(e0, e3)) & (e3 = op(e3, e0))) | ~ sP53), inference(usedef, [], [e107])).
fof(e107, plain, (sP53 <=> ((e0 = op(e0, e2)) & (e0 = op(e2, e0)) & (e3 = op(e0, e3)) & (e3 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP53])])).
fof(f108, plain, (((e0 = op(e0, e3)) & (e0 = op(e3, e0)) & (e3 = op(e0, e3)) & (e3 = op(e3, e0))) | ~ sP54), inference(usedef, [], [e108])).
fof(e108, plain, (sP54 <=> ((e0 = op(e0, e3)) & (e0 = op(e3, e0)) & (e3 = op(e0, e3)) & (e3 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP54])])).
fof(f109, plain, (((e1 = op(e1, e0)) & (e1 = op(e0, e1)) & (e3 = op(e1, e3)) & (e3 = op(e3, e1))) | ~ sP55), inference(usedef, [], [e109])).
fof(e109, plain, (sP55 <=> ((e1 = op(e1, e0)) & (e1 = op(e0, e1)) & (e3 = op(e1, e3)) & (e3 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP55])])).
fof(f110, plain, (((e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e3 = op(e1, e3)) & (e3 = op(e3, e1))) | ~ sP56), inference(usedef, [], [e110])).
fof(e110, plain, (sP56 <=> ((e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e3 = op(e1, e3)) & (e3 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP56])])).
fof(f111, plain, (((e1 = op(e1, e2)) & (e1 = op(e2, e1)) & (e3 = op(e1, e3)) & (e3 = op(e3, e1))) | ~ sP57), inference(usedef, [], [e111])).
fof(e111, plain, (sP57 <=> ((e1 = op(e1, e2)) & (e1 = op(e2, e1)) & (e3 = op(e1, e3)) & (e3 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP57])])).
fof(f112, plain, (((e1 = op(e1, e3)) & (e1 = op(e3, e1)) & (e3 = op(e1, e3)) & (e3 = op(e3, e1))) | ~ sP58), inference(usedef, [], [e112])).
fof(e112, plain, (sP58 <=> ((e1 = op(e1, e3)) & (e1 = op(e3, e1)) & (e3 = op(e1, e3)) & (e3 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP58])])).
fof(f113, plain, (((e2 = op(e2, e0)) & (e2 = op(e0, e2)) & (e3 = op(e2, e3)) & (e3 = op(e3, e2))) | ~ sP59), inference(usedef, [], [e113])).
fof(e113, plain, (sP59 <=> ((e2 = op(e2, e0)) & (e2 = op(e0, e2)) & (e3 = op(e2, e3)) & (e3 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP59])])).
fof(f114, plain, (((e2 = op(e2, e1)) & (e2 = op(e1, e2)) & (e3 = op(e2, e3)) & (e3 = op(e3, e2))) | ~ sP60), inference(usedef, [], [e114])).
fof(e114, plain, (sP60 <=> ((e2 = op(e2, e1)) & (e2 = op(e1, e2)) & (e3 = op(e2, e3)) & (e3 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP60])])).
fof(f115, plain, (((e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e3 = op(e2, e3)) & (e3 = op(e3, e2))) | ~ sP61), inference(usedef, [], [e115])).
fof(e115, plain, (sP61 <=> ((e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e3 = op(e2, e3)) & (e3 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP61])])).
fof(f116, plain, (((e2 = op(e2, e3)) & (e2 = op(e3, e2)) & (e3 = op(e2, e3)) & (e3 = op(e3, e2))) | ~ sP62), inference(usedef, [], [e116])).
fof(e116, plain, (sP62 <=> ((e2 = op(e2, e3)) & (e2 = op(e3, e2)) & (e3 = op(e2, e3)) & (e3 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP62])])).
fof(f117, plain, (((e3 = op(e3, e0)) & (e3 = op(e0, e3)) & (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ~ sP63), inference(usedef, [], [e117])).
fof(e117, plain, (sP63 <=> ((e3 = op(e3, e0)) & (e3 = op(e0, e3)) & (e3 = op(e3, e3)) & (e3 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP63])])).
fof(f118, plain, (((e3 = op(e3, e1)) & (e3 = op(e1, e3)) & (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ~ sP64), inference(usedef, [], [e118])).
fof(e118, plain, (sP64 <=> ((e3 = op(e3, e1)) & (e3 = op(e1, e3)) & (e3 = op(e3, e3)) & (e3 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP64])])).
fof(f119, plain, (((e3 = op(e3, e2)) & (e3 = op(e2, e3)) & (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ~ sP65), inference(usedef, [], [e119])).
fof(e119, plain, (sP65 <=> ((e3 = op(e3, e2)) & (e3 = op(e2, e3)) & (e3 = op(e3, e3)) & (e3 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP65])])).
fof(f5, plain, ((((e3 = op(e3, e3)) & (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ((e3 = op(e3, e2)) & (e3 = op(e2, e3)) & (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ((e3 = op(e3, e1)) & (e3 = op(e1, e3)) & (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ((e3 = op(e3, e0)) & (e3 = op(e0, e3)) & (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ((e2 = op(e2, e3)) & (e2 = op(e3, e2)) & (e3 = op(e2, e3)) & (e3 = op(e3, e2))) | ((e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e3 = op(e2, e3)) & (e3 = op(e3, e2))) | ((e2 = op(e2, e1)) & (e2 = op(e1, e2)) & (e3 = op(e2, e3)) & (e3 = op(e3, e2))) | ((e2 = op(e2, e0)) & (e2 = op(e0, e2)) & (e3 = op(e2, e3)) & (e3 = op(e3, e2))) | ((e1 = op(e1, e3)) & (e1 = op(e3, e1)) & (e3 = op(e1, e3)) & (e3 = op(e3, e1))) | ((e1 = op(e1, e2)) & (e1 = op(e2, e1)) & (e3 = op(e1, e3)) & (e3 = op(e3, e1))) | ((e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e3 = op(e1, e3)) & (e3 = op(e3, e1))) | ((e1 = op(e1, e0)) & (e1 = op(e0, e1)) & (e3 = op(e1, e3)) & (e3 = op(e3, e1))) | ((e0 = op(e0, e3)) & (e0 = op(e3, e0)) & (e3 = op(e0, e3)) & (e3 = op(e3, e0))) | ((e0 = op(e0, e2)) & (e0 = op(e2, e0)) & (e3 = op(e0, e3)) & (e3 = op(e3, e0))) | ((e0 = op(e0, e1)) & (e0 = op(e1, e0)) & (e3 = op(e0, e3)) & (e3 = op(e3, e0))) | ((e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e3 = op(e0, e3)) & (e3 = op(e3, e0))) | ((e3 = op(e3, e3)) & (e3 = op(e3, e3)) & (e2 = op(e3, e2)) & (e2 = op(e2, e3))) | ((e3 = op(e3, e2)) & (e3 = op(e2, e3)) & (e2 = op(e3, e2)) & (e2 = op(e2, e3))) | ((e3 = op(e3, e1)) & (e3 = op(e1, e3)) & (e2 = op(e3, e2)) & (e2 = op(e2, e3))) | ((e3 = op(e3, e0)) & (e3 = op(e0, e3)) & (e2 = op(e3, e2)) & (e2 = op(e2, e3))) | ((e2 = op(e2, e3)) & (e2 = op(e3, e2)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ((e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ((e2 = op(e2, e1)) & (e2 = op(e1, e2)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ((e2 = op(e2, e0)) & (e2 = op(e0, e2)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ((e1 = op(e1, e3)) & (e1 = op(e3, e1)) & (e2 = op(e1, e2)) & (e2 = op(e2, e1))) | ((e1 = op(e1, e2)) & (e1 = op(e2, e1)) & (e2 = op(e1, e2)) & (e2 = op(e2, e1))) | ((e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e2 = op(e1, e2)) & (e2 = op(e2, e1))) | ((e1 = op(e1, e0)) & (e1 = op(e0, e1)) & (e2 = op(e1, e2)) & (e2 = op(e2, e1))) | ((e0 = op(e0, e3)) & (e0 = op(e3, e0)) & (e2 = op(e0, e2)) & (e2 = op(e2, e0))) | ((e0 = op(e0, e2)) & (e0 = op(e2, e0)) & (e2 = op(e0, e2)) & (e2 = op(e2, e0))) | ((e0 = op(e0, e1)) & (e0 = op(e1, e0)) & (e2 = op(e0, e2)) & (e2 = op(e2, e0))) | ((e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e2 = op(e0, e2)) & (e2 = op(e2, e0))) | ((e3 = op(e3, e3)) & (e3 = op(e3, e3)) & (e1 = op(e3, e1)) & (e1 = op(e1, e3))) | ((e3 = op(e3, e2)) & (e3 = op(e2, e3)) & (e1 = op(e3, e1)) & (e1 = op(e1, e3))) | ((e3 = op(e3, e1)) & (e3 = op(e1, e3)) & (e1 = op(e3, e1)) & (e1 = op(e1, e3))) | ((e3 = op(e3, e0)) & (e3 = op(e0, e3)) & (e1 = op(e3, e1)) & (e1 = op(e1, e3))) | ((e2 = op(e2, e3)) & (e2 = op(e3, e2)) & (e1 = op(e2, e1)) & (e1 = op(e1, e2))) | ((e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e1 = op(e2, e1)) & (e1 = op(e1, e2))) | ((e2 = op(e2, e1)) & (e2 = op(e1, e2)) & (e1 = op(e2, e1)) & (e1 = op(e1, e2))) | ((e2 = op(e2, e0)) & (e2 = op(e0, e2)) & (e1 = op(e2, e1)) & (e1 = op(e1, e2))) | ((e1 = op(e1, e3)) & (e1 = op(e3, e1)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ((e1 = op(e1, e2)) & (e1 = op(e2, e1)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ((e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ((e1 = op(e1, e0)) & (e1 = op(e0, e1)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ((e0 = op(e0, e3)) & (e0 = op(e3, e0)) & (e1 = op(e0, e1)) & (e1 = op(e1, e0))) | ((e0 = op(e0, e2)) & (e0 = op(e2, e0)) & (e1 = op(e0, e1)) & (e1 = op(e1, e0))) | ((e0 = op(e0, e1)) & (e0 = op(e1, e0)) & (e1 = op(e0, e1)) & (e1 = op(e1, e0))) | ((e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e1 = op(e0, e1)) & (e1 = op(e1, e0))) | ((e3 = op(e3, e3)) & (e3 = op(e3, e3)) & (e0 = op(e3, e0)) & (e0 = op(e0, e3))) | ((e3 = op(e3, e2)) & (e3 = op(e2, e3)) & (e0 = op(e3, e0)) & (e0 = op(e0, e3))) | ((e3 = op(e3, e1)) & (e3 = op(e1, e3)) & (e0 = op(e3, e0)) & (e0 = op(e0, e3))) | ((e3 = op(e3, e0)) & (e3 = op(e0, e3)) & (e0 = op(e3, e0)) & (e0 = op(e0, e3))) | ((e2 = op(e2, e3)) & (e2 = op(e3, e2)) & (e0 = op(e2, e0)) & (e0 = op(e0, e2))) | ((e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e0 = op(e2, e0)) & (e0 = op(e0, e2))) | ((e2 = op(e2, e1)) & (e2 = op(e1, e2)) & (e0 = op(e2, e0)) & (e0 = op(e0, e2))) | ((e2 = op(e2, e0)) & (e2 = op(e0, e2)) & (e0 = op(e2, e0)) & (e0 = op(e0, e2))) | ((e1 = op(e1, e3)) & (e1 = op(e3, e1)) & (e0 = op(e1, e0)) & (e0 = op(e0, e1))) | ((e1 = op(e1, e2)) & (e1 = op(e2, e1)) & (e0 = op(e1, e0)) & (e0 = op(e0, e1))) | ((e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e0 = op(e1, e0)) & (e0 = op(e0, e1))) | ((e1 = op(e1, e0)) & (e1 = op(e0, e1)) & (e0 = op(e1, e0)) & (e0 = op(e0, e1))) | ((e0 = op(e0, e3)) & (e0 = op(e3, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ((e0 = op(e0, e2)) & (e0 = op(e2, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ((e0 = op(e0, e1)) & (e0 = op(e1, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ((e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0)))) & ((~ (e3 = op(e3, op(e3, e3))) & ~ (e2 = op(e3, op(e3, e2))) & ~ (e1 = op(e3, op(e3, e1))) & ~ (e0 = op(e3, op(e3, e0)))) | (~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0)))) | (~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0)))) | (~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))))) & ((e3 = op(e3, e3)) | (e3 = op(e2, e2)) | (e3 = op(e1, e1)) | (op(e0, e0) = e3)) & ((e2 = op(e3, e3)) | (e2 = op(e2, e2)) | (e2 = op(e1, e1)) | (op(e0, e0) = e2)) & ((e1 = op(e3, e3)) | (e1 = op(e2, e2)) | (e1 = op(e1, e1)) | (op(e0, e0) = e1)) & ((e0 = op(e3, e3)) | (e0 = op(e2, e2)) | (e0 = op(e1, e1)) | (e0 = op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG140+1.p', ax5)).
fof(f1495, plain, (spl66_62 | spl66_42 | spl66_22 | spl66_2), inference(avatar_split_clause, [], [f554, f594, f679, f764, f849])).
fof(f554, plain, ((e1 = op(e3, e3)) | (e1 = op(e2, e2)) | (e1 = op(e1, e1)) | (op(e0, e0) = e1)), inference(cnf_transformation, [], [f120])).
fof(f1494, plain, (spl66_63 | spl66_43 | spl66_23 | spl66_3), inference(avatar_split_clause, [], [f555, f598, f683, f768, f853])).
fof(f555, plain, ((e2 = op(e3, e3)) | (e2 = op(e2, e2)) | (e2 = op(e1, e1)) | (op(e0, e0) = e2)), inference(cnf_transformation, [], [f120])).
fof(f1493, plain, (spl66_64 | spl66_44 | spl66_24 | spl66_4), inference(avatar_split_clause, [], [f556, f602, f687, f772, f857])).
fof(f556, plain, ((e3 = op(e3, e3)) | (e3 = op(e2, e2)) | (e3 = op(e1, e1)) | (op(e0, e0) = e3)), inference(cnf_transformation, [], [f120])).
fof(f1492, plain, (spl66_138 | spl66_133 | spl66_128 | ~ spl66_146), inference(avatar_split_clause, [], [f557, f1489, f1398, f1422, f1446])).
fof(f1446, plain, (spl66_138 <=> sP0), introduced(avatar_definition, [new_symbols(naming, [spl66_138])])).
fof(f1422, plain, (spl66_133 <=> sP1), introduced(avatar_definition, [new_symbols(naming, [spl66_133])])).
fof(f1398, plain, (spl66_128 <=> sP2), introduced(avatar_definition, [new_symbols(naming, [spl66_128])])).
fof(f557, plain, (~ (e0 = op(e3, op(e3, e0))) | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f120])).
fof(f1487, plain, (spl66_138 | spl66_133 | spl66_128 | ~ spl66_145), inference(avatar_split_clause, [], [f558, f1484, f1398, f1422, f1446])).
fof(f558, plain, (~ (e1 = op(e3, op(e3, e1))) | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f120])).
fof(f1482, plain, (spl66_138 | spl66_133 | spl66_128 | ~ spl66_144), inference(avatar_split_clause, [], [f559, f1479, f1398, f1422, f1446])).
fof(f559, plain, (~ (e2 = op(e3, op(e3, e2))) | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f120])).
fof(f1477, plain, (spl66_138 | spl66_133 | spl66_128 | ~ spl66_143), inference(avatar_split_clause, [], [f560, f1474, f1398, f1422, f1446])).
fof(f560, plain, (~ (e3 = op(e3, op(e3, e3))) | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f120])).
fof(f1472, plain, (spl66_127 | spl66_126 | spl66_125 | spl66_124 | spl66_123 | spl66_122 | spl66_121 | spl66_120 | spl66_119 | spl66_118 | spl66_117 | spl66_116 | spl66_115 | spl66_114 | spl66_113 | spl66_112 | spl66_111 | spl66_110 | spl66_109 | spl66_108 | spl66_107 | spl66_106 | spl66_105 | spl66_104 | spl66_103 | spl66_102 | spl66_101 | spl66_100 | spl66_99 | spl66_98 | spl66_97 | spl66_96 | spl66_95 | spl66_94 | spl66_93 | spl66_92 | spl66_91 | spl66_90 | spl66_89 | spl66_88 | spl66_87 | spl66_86 | spl66_85 | spl66_84 | spl66_83 | spl66_82 | spl66_81 | spl66_80 | spl66_79 | spl66_78 | spl66_77 | spl66_76 | spl66_75 | spl66_74 | spl66_73 | spl66_72 | spl66_71 | spl66_70 | spl66_69 | spl66_68 | spl66_67 | spl66_66 | spl66_65 | spl66_4), inference(avatar_split_clause, [], [f561, f602, f894, f902, f910, f918, f926, f934, f942, f950, f958, f966, f974, f982, f990, f998, f1006, f1014, f1022, f1030, f1038, f1046, f1054, f1062, f1070, f1078, f1086, f1094, f1102, f1110, f1118, f1126, f1134, f1142, f1150, f1158, f1166, f1174, f1182, f1190, f1198, f1206, f1214, f1222, f1230, f1238, f1246, f1254, f1262, f1270, f1278, f1286, f1294, f1302, f1310, f1318, f1326, f1334, f1342, f1350, f1358, f1366, f1374, f1382, f1390])).
fof(f1390, plain, (spl66_127 <=> sP3), introduced(avatar_definition, [new_symbols(naming, [spl66_127])])).
fof(f1382, plain, (spl66_126 <=> sP4), introduced(avatar_definition, [new_symbols(naming, [spl66_126])])).
fof(f1374, plain, (spl66_125 <=> sP5), introduced(avatar_definition, [new_symbols(naming, [spl66_125])])).
fof(f1366, plain, (spl66_124 <=> sP6), introduced(avatar_definition, [new_symbols(naming, [spl66_124])])).
fof(f1358, plain, (spl66_123 <=> sP7), introduced(avatar_definition, [new_symbols(naming, [spl66_123])])).
fof(f1350, plain, (spl66_122 <=> sP8), introduced(avatar_definition, [new_symbols(naming, [spl66_122])])).
fof(f1342, plain, (spl66_121 <=> sP9), introduced(avatar_definition, [new_symbols(naming, [spl66_121])])).
fof(f1334, plain, (spl66_120 <=> sP10), introduced(avatar_definition, [new_symbols(naming, [spl66_120])])).
fof(f1326, plain, (spl66_119 <=> sP11), introduced(avatar_definition, [new_symbols(naming, [spl66_119])])).
fof(f1318, plain, (spl66_118 <=> sP12), introduced(avatar_definition, [new_symbols(naming, [spl66_118])])).
fof(f1310, plain, (spl66_117 <=> sP13), introduced(avatar_definition, [new_symbols(naming, [spl66_117])])).
fof(f1302, plain, (spl66_116 <=> sP14), introduced(avatar_definition, [new_symbols(naming, [spl66_116])])).
fof(f1294, plain, (spl66_115 <=> sP15), introduced(avatar_definition, [new_symbols(naming, [spl66_115])])).
fof(f1286, plain, (spl66_114 <=> sP16), introduced(avatar_definition, [new_symbols(naming, [spl66_114])])).
fof(f1278, plain, (spl66_113 <=> sP17), introduced(avatar_definition, [new_symbols(naming, [spl66_113])])).
fof(f1270, plain, (spl66_112 <=> sP18), introduced(avatar_definition, [new_symbols(naming, [spl66_112])])).
fof(f1262, plain, (spl66_111 <=> sP19), introduced(avatar_definition, [new_symbols(naming, [spl66_111])])).
fof(f1254, plain, (spl66_110 <=> sP20), introduced(avatar_definition, [new_symbols(naming, [spl66_110])])).
fof(f1246, plain, (spl66_109 <=> sP21), introduced(avatar_definition, [new_symbols(naming, [spl66_109])])).
fof(f1238, plain, (spl66_108 <=> sP22), introduced(avatar_definition, [new_symbols(naming, [spl66_108])])).
fof(f1230, plain, (spl66_107 <=> sP23), introduced(avatar_definition, [new_symbols(naming, [spl66_107])])).
fof(f1222, plain, (spl66_106 <=> sP24), introduced(avatar_definition, [new_symbols(naming, [spl66_106])])).
fof(f1214, plain, (spl66_105 <=> sP25), introduced(avatar_definition, [new_symbols(naming, [spl66_105])])).
fof(f1206, plain, (spl66_104 <=> sP26), introduced(avatar_definition, [new_symbols(naming, [spl66_104])])).
fof(f1198, plain, (spl66_103 <=> sP27), introduced(avatar_definition, [new_symbols(naming, [spl66_103])])).
fof(f1190, plain, (spl66_102 <=> sP28), introduced(avatar_definition, [new_symbols(naming, [spl66_102])])).
fof(f1182, plain, (spl66_101 <=> sP29), introduced(avatar_definition, [new_symbols(naming, [spl66_101])])).
fof(f1174, plain, (spl66_100 <=> sP30), introduced(avatar_definition, [new_symbols(naming, [spl66_100])])).
fof(f1166, plain, (spl66_99 <=> sP31), introduced(avatar_definition, [new_symbols(naming, [spl66_99])])).
fof(f1158, plain, (spl66_98 <=> sP32), introduced(avatar_definition, [new_symbols(naming, [spl66_98])])).
fof(f1150, plain, (spl66_97 <=> sP33), introduced(avatar_definition, [new_symbols(naming, [spl66_97])])).
fof(f1142, plain, (spl66_96 <=> sP34), introduced(avatar_definition, [new_symbols(naming, [spl66_96])])).
fof(f1134, plain, (spl66_95 <=> sP35), introduced(avatar_definition, [new_symbols(naming, [spl66_95])])).
fof(f1126, plain, (spl66_94 <=> sP36), introduced(avatar_definition, [new_symbols(naming, [spl66_94])])).
fof(f1118, plain, (spl66_93 <=> sP37), introduced(avatar_definition, [new_symbols(naming, [spl66_93])])).
fof(f1110, plain, (spl66_92 <=> sP38), introduced(avatar_definition, [new_symbols(naming, [spl66_92])])).
fof(f1102, plain, (spl66_91 <=> sP39), introduced(avatar_definition, [new_symbols(naming, [spl66_91])])).
fof(f1094, plain, (spl66_90 <=> sP40), introduced(avatar_definition, [new_symbols(naming, [spl66_90])])).
fof(f1086, plain, (spl66_89 <=> sP41), introduced(avatar_definition, [new_symbols(naming, [spl66_89])])).
fof(f1078, plain, (spl66_88 <=> sP42), introduced(avatar_definition, [new_symbols(naming, [spl66_88])])).
fof(f1070, plain, (spl66_87 <=> sP43), introduced(avatar_definition, [new_symbols(naming, [spl66_87])])).
fof(f1062, plain, (spl66_86 <=> sP44), introduced(avatar_definition, [new_symbols(naming, [spl66_86])])).
fof(f1054, plain, (spl66_85 <=> sP45), introduced(avatar_definition, [new_symbols(naming, [spl66_85])])).
fof(f1046, plain, (spl66_84 <=> sP46), introduced(avatar_definition, [new_symbols(naming, [spl66_84])])).
fof(f1038, plain, (spl66_83 <=> sP47), introduced(avatar_definition, [new_symbols(naming, [spl66_83])])).
fof(f1030, plain, (spl66_82 <=> sP48), introduced(avatar_definition, [new_symbols(naming, [spl66_82])])).
fof(f1022, plain, (spl66_81 <=> sP49), introduced(avatar_definition, [new_symbols(naming, [spl66_81])])).
fof(f1014, plain, (spl66_80 <=> sP50), introduced(avatar_definition, [new_symbols(naming, [spl66_80])])).
fof(f1006, plain, (spl66_79 <=> sP51), introduced(avatar_definition, [new_symbols(naming, [spl66_79])])).
fof(f998, plain, (spl66_78 <=> sP52), introduced(avatar_definition, [new_symbols(naming, [spl66_78])])).
fof(f990, plain, (spl66_77 <=> sP53), introduced(avatar_definition, [new_symbols(naming, [spl66_77])])).
fof(f982, plain, (spl66_76 <=> sP54), introduced(avatar_definition, [new_symbols(naming, [spl66_76])])).
fof(f974, plain, (spl66_75 <=> sP55), introduced(avatar_definition, [new_symbols(naming, [spl66_75])])).
fof(f966, plain, (spl66_74 <=> sP56), introduced(avatar_definition, [new_symbols(naming, [spl66_74])])).
fof(f958, plain, (spl66_73 <=> sP57), introduced(avatar_definition, [new_symbols(naming, [spl66_73])])).
fof(f950, plain, (spl66_72 <=> sP58), introduced(avatar_definition, [new_symbols(naming, [spl66_72])])).
fof(f942, plain, (spl66_71 <=> sP59), introduced(avatar_definition, [new_symbols(naming, [spl66_71])])).
fof(f934, plain, (spl66_70 <=> sP60), introduced(avatar_definition, [new_symbols(naming, [spl66_70])])).
fof(f926, plain, (spl66_69 <=> sP61), introduced(avatar_definition, [new_symbols(naming, [spl66_69])])).
fof(f918, plain, (spl66_68 <=> sP62), introduced(avatar_definition, [new_symbols(naming, [spl66_68])])).
fof(f910, plain, (spl66_67 <=> sP63), introduced(avatar_definition, [new_symbols(naming, [spl66_67])])).
fof(f902, plain, (spl66_66 <=> sP64), introduced(avatar_definition, [new_symbols(naming, [spl66_66])])).
fof(f894, plain, (spl66_65 <=> sP65), introduced(avatar_definition, [new_symbols(naming, [spl66_65])])).
fof(f561, plain, ((e3 = op(e3, e3)) | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3), inference(cnf_transformation, [], [f120])).
fof(f1468, plain, (~ spl66_138 | ~ spl66_142), inference(avatar_split_clause, [], [f549, f1465, f1446])).
fof(f549, plain, (~ (e0 = op(e0, op(e0, e0))) | ~ sP0), inference(cnf_transformation, [], [f186])).
fof(f186, plain, ((~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0)))) | ~ sP0), inference(nnf_transformation, [], [f54])).
fof(f1463, plain, (~ spl66_138 | ~ spl66_141), inference(avatar_split_clause, [], [f550, f1460, f1446])).
fof(f550, plain, (~ (e1 = op(e0, op(e0, e1))) | ~ sP0), inference(cnf_transformation, [], [f186])).
fof(f1458, plain, (~ spl66_138 | ~ spl66_140), inference(avatar_split_clause, [], [f551, f1455, f1446])).
fof(f551, plain, (~ (e2 = op(e0, op(e0, e2))) | ~ sP0), inference(cnf_transformation, [], [f186])).
fof(f1453, plain, (~ spl66_138 | ~ spl66_139), inference(avatar_split_clause, [], [f552, f1450, f1446])).
fof(f552, plain, (~ (e3 = op(e0, op(e0, e3))) | ~ sP0), inference(cnf_transformation, [], [f186])).
fof(f1444, plain, (~ spl66_133 | ~ spl66_137), inference(avatar_split_clause, [], [f545, f1441, f1422])).
fof(f545, plain, (~ (e0 = op(e1, op(e1, e0))) | ~ sP1), inference(cnf_transformation, [], [f185])).
fof(f185, plain, ((~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0)))) | ~ sP1), inference(nnf_transformation, [], [f55])).
fof(f1439, plain, (~ spl66_133 | ~ spl66_136), inference(avatar_split_clause, [], [f546, f1436, f1422])).
fof(f546, plain, (~ (e1 = op(e1, op(e1, e1))) | ~ sP1), inference(cnf_transformation, [], [f185])).
fof(f1434, plain, (~ spl66_133 | ~ spl66_135), inference(avatar_split_clause, [], [f547, f1431, f1422])).
fof(f547, plain, (~ (e2 = op(e1, op(e1, e2))) | ~ sP1), inference(cnf_transformation, [], [f185])).
fof(f1429, plain, (~ spl66_133 | ~ spl66_134), inference(avatar_split_clause, [], [f548, f1426, f1422])).
fof(f548, plain, (~ (e3 = op(e1, op(e1, e3))) | ~ sP1), inference(cnf_transformation, [], [f185])).
fof(f1420, plain, (~ spl66_128 | ~ spl66_132), inference(avatar_split_clause, [], [f541, f1417, f1398])).
fof(f541, plain, (~ (e0 = op(e2, op(e2, e0))) | ~ sP2), inference(cnf_transformation, [], [f184])).
fof(f184, plain, ((~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0)))) | ~ sP2), inference(nnf_transformation, [], [f56])).
fof(f1415, plain, (~ spl66_128 | ~ spl66_131), inference(avatar_split_clause, [], [f542, f1412, f1398])).
fof(f542, plain, (~ (e1 = op(e2, op(e2, e1))) | ~ sP2), inference(cnf_transformation, [], [f184])).
fof(f1410, plain, (~ spl66_128 | ~ spl66_130), inference(avatar_split_clause, [], [f543, f1407, f1398])).
fof(f543, plain, (~ (e2 = op(e2, op(e2, e2))) | ~ sP2), inference(cnf_transformation, [], [f184])).
fof(f1405, plain, (~ spl66_128 | ~ spl66_129), inference(avatar_split_clause, [], [f544, f1402, f1398])).
fof(f544, plain, (~ (e3 = op(e2, op(e2, e3))) | ~ sP2), inference(cnf_transformation, [], [f184])).
fof(f1396, plain, (~ spl66_127 | spl66_61), inference(avatar_split_clause, [], [f537, f845, f1390])).
fof(f537, plain, ((e0 = op(e0, e0)) | ~ sP3), inference(cnf_transformation, [], [f183])).
fof(f183, plain, (((e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP3), inference(nnf_transformation, [], [f57])).
fof(f1388, plain, (~ spl66_126 | spl66_61), inference(avatar_split_clause, [], [f533, f845, f1382])).
fof(f533, plain, ((e0 = op(e0, e0)) | ~ sP4), inference(cnf_transformation, [], [f182])).
fof(f182, plain, (((e0 = op(e0, e1)) & (e0 = op(e1, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP4), inference(nnf_transformation, [], [f58])).
fof(f1386, plain, (~ spl66_126 | spl66_45), inference(avatar_split_clause, [], [f535, f777, f1382])).
fof(f535, plain, ((e0 = op(e1, e0)) | ~ sP4), inference(cnf_transformation, [], [f182])).
fof(f1385, plain, (~ spl66_126 | spl66_57), inference(avatar_split_clause, [], [f536, f828, f1382])).
fof(f536, plain, ((e0 = op(e0, e1)) | ~ sP4), inference(cnf_transformation, [], [f182])).
fof(f1380, plain, (~ spl66_125 | spl66_61), inference(avatar_split_clause, [], [f529, f845, f1374])).
fof(f529, plain, ((e0 = op(e0, e0)) | ~ sP5), inference(cnf_transformation, [], [f181])).
fof(f181, plain, (((e0 = op(e0, e2)) & (e0 = op(e2, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP5), inference(nnf_transformation, [], [f59])).
fof(f1378, plain, (~ spl66_125 | spl66_29), inference(avatar_split_clause, [], [f531, f709, f1374])).
fof(f531, plain, ((e0 = op(e2, e0)) | ~ sP5), inference(cnf_transformation, [], [f181])).
fof(f1377, plain, (~ spl66_125 | spl66_53), inference(avatar_split_clause, [], [f532, f811, f1374])).
fof(f532, plain, ((e0 = op(e0, e2)) | ~ sP5), inference(cnf_transformation, [], [f181])).
fof(f1372, plain, (~ spl66_124 | spl66_61), inference(avatar_split_clause, [], [f525, f845, f1366])).
fof(f525, plain, ((e0 = op(e0, e0)) | ~ sP6), inference(cnf_transformation, [], [f180])).
fof(f180, plain, (((e0 = op(e0, e3)) & (e0 = op(e3, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP6), inference(nnf_transformation, [], [f60])).
fof(f1370, plain, (~ spl66_124 | spl66_13), inference(avatar_split_clause, [], [f527, f641, f1366])).
fof(f527, plain, ((e0 = op(e3, e0)) | ~ sP6), inference(cnf_transformation, [], [f180])).
fof(f1369, plain, (~ spl66_124 | spl66_49), inference(avatar_split_clause, [], [f528, f794, f1366])).
fof(f528, plain, ((e0 = op(e0, e3)) | ~ sP6), inference(cnf_transformation, [], [f180])).
fof(f1364, plain, (~ spl66_123 | spl66_57), inference(avatar_split_clause, [], [f521, f828, f1358])).
fof(f521, plain, ((e0 = op(e0, e1)) | ~ sP7), inference(cnf_transformation, [], [f179])).
fof(f179, plain, (((e1 = op(e1, e0)) & (e1 = op(e0, e1)) & (e0 = op(e1, e0)) & (e0 = op(e0, e1))) | ~ sP7), inference(nnf_transformation, [], [f61])).
fof(f1363, plain, (~ spl66_123 | spl66_45), inference(avatar_split_clause, [], [f522, f777, f1358])).
fof(f522, plain, ((e0 = op(e1, e0)) | ~ sP7), inference(cnf_transformation, [], [f179])).
fof(f1362, plain, (~ spl66_123 | spl66_58), inference(avatar_split_clause, [], [f523, f832, f1358])).
fof(f523, plain, ((e1 = op(e0, e1)) | ~ sP7), inference(cnf_transformation, [], [f179])).
fof(f1361, plain, (~ spl66_123 | spl66_46), inference(avatar_split_clause, [], [f524, f781, f1358])).
fof(f524, plain, ((e1 = op(e1, e0)) | ~ sP7), inference(cnf_transformation, [], [f179])).
fof(f1356, plain, (~ spl66_122 | spl66_57), inference(avatar_split_clause, [], [f517, f828, f1350])).
fof(f517, plain, ((e0 = op(e0, e1)) | ~ sP8), inference(cnf_transformation, [], [f178])).
fof(f178, plain, (((e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e0 = op(e1, e0)) & (e0 = op(e0, e1))) | ~ sP8), inference(nnf_transformation, [], [f62])).
fof(f1355, plain, (~ spl66_122 | spl66_45), inference(avatar_split_clause, [], [f518, f777, f1350])).
fof(f518, plain, ((e0 = op(e1, e0)) | ~ sP8), inference(cnf_transformation, [], [f178])).
fof(f1354, plain, (~ spl66_122 | spl66_42), inference(avatar_split_clause, [], [f519, f764, f1350])).
fof(f519, plain, ((e1 = op(e1, e1)) | ~ sP8), inference(cnf_transformation, [], [f178])).
fof(f1348, plain, (~ spl66_121 | spl66_57), inference(avatar_split_clause, [], [f513, f828, f1342])).
fof(f513, plain, ((e0 = op(e0, e1)) | ~ sP9), inference(cnf_transformation, [], [f177])).
fof(f177, plain, (((e1 = op(e1, e2)) & (e1 = op(e2, e1)) & (e0 = op(e1, e0)) & (e0 = op(e0, e1))) | ~ sP9), inference(nnf_transformation, [], [f63])).
fof(f1347, plain, (~ spl66_121 | spl66_45), inference(avatar_split_clause, [], [f514, f777, f1342])).
fof(f514, plain, ((e0 = op(e1, e0)) | ~ sP9), inference(cnf_transformation, [], [f177])).
fof(f1345, plain, (~ spl66_121 | spl66_38), inference(avatar_split_clause, [], [f516, f747, f1342])).
fof(f516, plain, ((e1 = op(e1, e2)) | ~ sP9), inference(cnf_transformation, [], [f177])).
fof(f1340, plain, (~ spl66_120 | spl66_57), inference(avatar_split_clause, [], [f509, f828, f1334])).
fof(f509, plain, ((e0 = op(e0, e1)) | ~ sP10), inference(cnf_transformation, [], [f176])).
fof(f176, plain, (((e1 = op(e1, e3)) & (e1 = op(e3, e1)) & (e0 = op(e1, e0)) & (e0 = op(e0, e1))) | ~ sP10), inference(nnf_transformation, [], [f64])).
fof(f1339, plain, (~ spl66_120 | spl66_45), inference(avatar_split_clause, [], [f510, f777, f1334])).
fof(f510, plain, ((e0 = op(e1, e0)) | ~ sP10), inference(cnf_transformation, [], [f176])).
fof(f1338, plain, (~ spl66_120 | spl66_10), inference(avatar_split_clause, [], [f511, f628, f1334])).
fof(f511, plain, ((e1 = op(e3, e1)) | ~ sP10), inference(cnf_transformation, [], [f176])).
fof(f1337, plain, (~ spl66_120 | spl66_34), inference(avatar_split_clause, [], [f512, f730, f1334])).
fof(f512, plain, ((e1 = op(e1, e3)) | ~ sP10), inference(cnf_transformation, [], [f176])).
fof(f1332, plain, (~ spl66_119 | spl66_53), inference(avatar_split_clause, [], [f505, f811, f1326])).
fof(f505, plain, ((e0 = op(e0, e2)) | ~ sP11), inference(cnf_transformation, [], [f175])).
fof(f175, plain, (((e2 = op(e2, e0)) & (e2 = op(e0, e2)) & (e0 = op(e2, e0)) & (e0 = op(e0, e2))) | ~ sP11), inference(nnf_transformation, [], [f65])).
fof(f1330, plain, (~ spl66_119 | spl66_55), inference(avatar_split_clause, [], [f507, f819, f1326])).
fof(f507, plain, ((e2 = op(e0, e2)) | ~ sP11), inference(cnf_transformation, [], [f175])).
fof(f1329, plain, (~ spl66_119 | spl66_31), inference(avatar_split_clause, [], [f508, f717, f1326])).
fof(f508, plain, ((e2 = op(e2, e0)) | ~ sP11), inference(cnf_transformation, [], [f175])).
fof(f1324, plain, (~ spl66_118 | spl66_53), inference(avatar_split_clause, [], [f501, f811, f1318])).
fof(f501, plain, ((e0 = op(e0, e2)) | ~ sP12), inference(cnf_transformation, [], [f174])).
fof(f174, plain, (((e2 = op(e2, e1)) & (e2 = op(e1, e2)) & (e0 = op(e2, e0)) & (e0 = op(e0, e2))) | ~ sP12), inference(nnf_transformation, [], [f66])).
fof(f1322, plain, (~ spl66_118 | spl66_39), inference(avatar_split_clause, [], [f503, f751, f1318])).
fof(f503, plain, ((e2 = op(e1, e2)) | ~ sP12), inference(cnf_transformation, [], [f174])).
fof(f1321, plain, (~ spl66_118 | spl66_27), inference(avatar_split_clause, [], [f504, f700, f1318])).
fof(f504, plain, ((e2 = op(e2, e1)) | ~ sP12), inference(cnf_transformation, [], [f174])).
fof(f1316, plain, (~ spl66_117 | spl66_53), inference(avatar_split_clause, [], [f497, f811, f1310])).
fof(f497, plain, ((e0 = op(e0, e2)) | ~ sP13), inference(cnf_transformation, [], [f173])).
fof(f173, plain, (((e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e0 = op(e2, e0)) & (e0 = op(e0, e2))) | ~ sP13), inference(nnf_transformation, [], [f67])).
fof(f1314, plain, (~ spl66_117 | spl66_23), inference(avatar_split_clause, [], [f499, f683, f1310])).
fof(f499, plain, ((e2 = op(e2, e2)) | ~ sP13), inference(cnf_transformation, [], [f173])).
fof(f1308, plain, (~ spl66_116 | spl66_53), inference(avatar_split_clause, [], [f493, f811, f1302])).
fof(f493, plain, ((e0 = op(e0, e2)) | ~ sP14), inference(cnf_transformation, [], [f172])).
fof(f172, plain, (((e2 = op(e2, e3)) & (e2 = op(e3, e2)) & (e0 = op(e2, e0)) & (e0 = op(e0, e2))) | ~ sP14), inference(nnf_transformation, [], [f68])).
fof(f1306, plain, (~ spl66_116 | spl66_7), inference(avatar_split_clause, [], [f495, f615, f1302])).
fof(f495, plain, ((e2 = op(e3, e2)) | ~ sP14), inference(cnf_transformation, [], [f172])).
fof(f1305, plain, (~ spl66_116 | spl66_19), inference(avatar_split_clause, [], [f496, f666, f1302])).
fof(f496, plain, ((e2 = op(e2, e3)) | ~ sP14), inference(cnf_transformation, [], [f172])).
fof(f1300, plain, (~ spl66_115 | spl66_49), inference(avatar_split_clause, [], [f489, f794, f1294])).
fof(f489, plain, ((e0 = op(e0, e3)) | ~ sP15), inference(cnf_transformation, [], [f171])).
fof(f171, plain, (((e3 = op(e3, e0)) & (e3 = op(e0, e3)) & (e0 = op(e3, e0)) & (e0 = op(e0, e3))) | ~ sP15), inference(nnf_transformation, [], [f69])).
fof(f1299, plain, (~ spl66_115 | spl66_13), inference(avatar_split_clause, [], [f490, f641, f1294])).
fof(f490, plain, ((e0 = op(e3, e0)) | ~ sP15), inference(cnf_transformation, [], [f171])).
fof(f1298, plain, (~ spl66_115 | spl66_52), inference(avatar_split_clause, [], [f491, f806, f1294])).
fof(f491, plain, ((e3 = op(e0, e3)) | ~ sP15), inference(cnf_transformation, [], [f171])).
fof(f1297, plain, (~ spl66_115 | spl66_16), inference(avatar_split_clause, [], [f492, f653, f1294])).
fof(f492, plain, ((e3 = op(e3, e0)) | ~ sP15), inference(cnf_transformation, [], [f171])).
fof(f1292, plain, (~ spl66_114 | spl66_49), inference(avatar_split_clause, [], [f485, f794, f1286])).
fof(f485, plain, ((e0 = op(e0, e3)) | ~ sP16), inference(cnf_transformation, [], [f170])).
fof(f170, plain, (((e3 = op(e3, e1)) & (e3 = op(e1, e3)) & (e0 = op(e3, e0)) & (e0 = op(e0, e3))) | ~ sP16), inference(nnf_transformation, [], [f70])).
fof(f1291, plain, (~ spl66_114 | spl66_13), inference(avatar_split_clause, [], [f486, f641, f1286])).
fof(f486, plain, ((e0 = op(e3, e0)) | ~ sP16), inference(cnf_transformation, [], [f170])).
fof(f1290, plain, (~ spl66_114 | spl66_36), inference(avatar_split_clause, [], [f487, f738, f1286])).
fof(f487, plain, ((e3 = op(e1, e3)) | ~ sP16), inference(cnf_transformation, [], [f170])).
fof(f1289, plain, (~ spl66_114 | spl66_12), inference(avatar_split_clause, [], [f488, f636, f1286])).
fof(f488, plain, ((e3 = op(e3, e1)) | ~ sP16), inference(cnf_transformation, [], [f170])).
fof(f1284, plain, (~ spl66_113 | spl66_49), inference(avatar_split_clause, [], [f481, f794, f1278])).
fof(f481, plain, ((e0 = op(e0, e3)) | ~ sP17), inference(cnf_transformation, [], [f169])).
fof(f169, plain, (((e3 = op(e3, e2)) & (e3 = op(e2, e3)) & (e0 = op(e3, e0)) & (e0 = op(e0, e3))) | ~ sP17), inference(nnf_transformation, [], [f71])).
fof(f1283, plain, (~ spl66_113 | spl66_13), inference(avatar_split_clause, [], [f482, f641, f1278])).
fof(f482, plain, ((e0 = op(e3, e0)) | ~ sP17), inference(cnf_transformation, [], [f169])).
fof(f1282, plain, (~ spl66_113 | spl66_20), inference(avatar_split_clause, [], [f483, f670, f1278])).
fof(f483, plain, ((e3 = op(e2, e3)) | ~ sP17), inference(cnf_transformation, [], [f169])).
fof(f1281, plain, (~ spl66_113 | spl66_8), inference(avatar_split_clause, [], [f484, f619, f1278])).
fof(f484, plain, ((e3 = op(e3, e2)) | ~ sP17), inference(cnf_transformation, [], [f169])).
fof(f1276, plain, (~ spl66_112 | spl66_49), inference(avatar_split_clause, [], [f477, f794, f1270])).
fof(f477, plain, ((e0 = op(e0, e3)) | ~ sP18), inference(cnf_transformation, [], [f168])).
fof(f168, plain, (((e3 = op(e3, e3)) & (e3 = op(e3, e3)) & (e0 = op(e3, e0)) & (e0 = op(e0, e3))) | ~ sP18), inference(nnf_transformation, [], [f72])).
fof(f1275, plain, (~ spl66_112 | spl66_13), inference(avatar_split_clause, [], [f478, f641, f1270])).
fof(f478, plain, ((e0 = op(e3, e0)) | ~ sP18), inference(cnf_transformation, [], [f168])).
fof(f1274, plain, (~ spl66_112 | spl66_4), inference(avatar_split_clause, [], [f479, f602, f1270])).
fof(f479, plain, ((e3 = op(e3, e3)) | ~ sP18), inference(cnf_transformation, [], [f168])).
fof(f1268, plain, (~ spl66_111 | spl66_46), inference(avatar_split_clause, [], [f473, f781, f1262])).
fof(f473, plain, ((e1 = op(e1, e0)) | ~ sP19), inference(cnf_transformation, [], [f167])).
fof(f167, plain, (((e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e1 = op(e0, e1)) & (e1 = op(e1, e0))) | ~ sP19), inference(nnf_transformation, [], [f73])).
fof(f1267, plain, (~ spl66_111 | spl66_58), inference(avatar_split_clause, [], [f474, f832, f1262])).
fof(f474, plain, ((e1 = op(e0, e1)) | ~ sP19), inference(cnf_transformation, [], [f167])).
fof(f1266, plain, (~ spl66_111 | spl66_61), inference(avatar_split_clause, [], [f475, f845, f1262])).
fof(f475, plain, ((e0 = op(e0, e0)) | ~ sP19), inference(cnf_transformation, [], [f167])).
fof(f1260, plain, (~ spl66_110 | spl66_46), inference(avatar_split_clause, [], [f469, f781, f1254])).
fof(f469, plain, ((e1 = op(e1, e0)) | ~ sP20), inference(cnf_transformation, [], [f166])).
fof(f166, plain, (((e0 = op(e0, e1)) & (e0 = op(e1, e0)) & (e1 = op(e0, e1)) & (e1 = op(e1, e0))) | ~ sP20), inference(nnf_transformation, [], [f74])).
fof(f1259, plain, (~ spl66_110 | spl66_58), inference(avatar_split_clause, [], [f470, f832, f1254])).
fof(f470, plain, ((e1 = op(e0, e1)) | ~ sP20), inference(cnf_transformation, [], [f166])).
fof(f1258, plain, (~ spl66_110 | spl66_45), inference(avatar_split_clause, [], [f471, f777, f1254])).
fof(f471, plain, ((e0 = op(e1, e0)) | ~ sP20), inference(cnf_transformation, [], [f166])).
fof(f1257, plain, (~ spl66_110 | spl66_57), inference(avatar_split_clause, [], [f472, f828, f1254])).
fof(f472, plain, ((e0 = op(e0, e1)) | ~ sP20), inference(cnf_transformation, [], [f166])).
fof(f1252, plain, (~ spl66_109 | spl66_46), inference(avatar_split_clause, [], [f465, f781, f1246])).
fof(f465, plain, ((e1 = op(e1, e0)) | ~ sP21), inference(cnf_transformation, [], [f165])).
fof(f165, plain, (((e0 = op(e0, e2)) & (e0 = op(e2, e0)) & (e1 = op(e0, e1)) & (e1 = op(e1, e0))) | ~ sP21), inference(nnf_transformation, [], [f75])).
fof(f1251, plain, (~ spl66_109 | spl66_58), inference(avatar_split_clause, [], [f466, f832, f1246])).
fof(f466, plain, ((e1 = op(e0, e1)) | ~ sP21), inference(cnf_transformation, [], [f165])).
fof(f1250, plain, (~ spl66_109 | spl66_29), inference(avatar_split_clause, [], [f467, f709, f1246])).
fof(f467, plain, ((e0 = op(e2, e0)) | ~ sP21), inference(cnf_transformation, [], [f165])).
fof(f1249, plain, (~ spl66_109 | spl66_53), inference(avatar_split_clause, [], [f468, f811, f1246])).
fof(f468, plain, ((e0 = op(e0, e2)) | ~ sP21), inference(cnf_transformation, [], [f165])).
fof(f1244, plain, (~ spl66_108 | spl66_46), inference(avatar_split_clause, [], [f461, f781, f1238])).
fof(f461, plain, ((e1 = op(e1, e0)) | ~ sP22), inference(cnf_transformation, [], [f164])).
fof(f164, plain, (((e0 = op(e0, e3)) & (e0 = op(e3, e0)) & (e1 = op(e0, e1)) & (e1 = op(e1, e0))) | ~ sP22), inference(nnf_transformation, [], [f76])).
fof(f1243, plain, (~ spl66_108 | spl66_58), inference(avatar_split_clause, [], [f462, f832, f1238])).
fof(f462, plain, ((e1 = op(e0, e1)) | ~ sP22), inference(cnf_transformation, [], [f164])).
fof(f1242, plain, (~ spl66_108 | spl66_13), inference(avatar_split_clause, [], [f463, f641, f1238])).
fof(f463, plain, ((e0 = op(e3, e0)) | ~ sP22), inference(cnf_transformation, [], [f164])).
fof(f1241, plain, (~ spl66_108 | spl66_49), inference(avatar_split_clause, [], [f464, f794, f1238])).
fof(f464, plain, ((e0 = op(e0, e3)) | ~ sP22), inference(cnf_transformation, [], [f164])).
fof(f1236, plain, (~ spl66_107 | spl66_42), inference(avatar_split_clause, [], [f457, f764, f1230])).
fof(f457, plain, ((e1 = op(e1, e1)) | ~ sP23), inference(cnf_transformation, [], [f163])).
fof(f163, plain, (((e1 = op(e1, e0)) & (e1 = op(e0, e1)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP23), inference(nnf_transformation, [], [f77])).
fof(f1234, plain, (~ spl66_107 | spl66_58), inference(avatar_split_clause, [], [f459, f832, f1230])).
fof(f459, plain, ((e1 = op(e0, e1)) | ~ sP23), inference(cnf_transformation, [], [f163])).
fof(f1233, plain, (~ spl66_107 | spl66_46), inference(avatar_split_clause, [], [f460, f781, f1230])).
fof(f460, plain, ((e1 = op(e1, e0)) | ~ sP23), inference(cnf_transformation, [], [f163])).
fof(f1228, plain, (~ spl66_106 | spl66_42), inference(avatar_split_clause, [], [f453, f764, f1222])).
fof(f453, plain, ((e1 = op(e1, e1)) | ~ sP24), inference(cnf_transformation, [], [f162])).
fof(f162, plain, (((e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP24), inference(nnf_transformation, [], [f78])).
fof(f1220, plain, (~ spl66_105 | spl66_42), inference(avatar_split_clause, [], [f449, f764, f1214])).
fof(f449, plain, ((e1 = op(e1, e1)) | ~ sP25), inference(cnf_transformation, [], [f161])).
fof(f161, plain, (((e1 = op(e1, e2)) & (e1 = op(e2, e1)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP25), inference(nnf_transformation, [], [f79])).
fof(f1218, plain, (~ spl66_105 | spl66_26), inference(avatar_split_clause, [], [f451, f696, f1214])).
fof(f451, plain, ((e1 = op(e2, e1)) | ~ sP25), inference(cnf_transformation, [], [f161])).
fof(f1217, plain, (~ spl66_105 | spl66_38), inference(avatar_split_clause, [], [f452, f747, f1214])).
fof(f452, plain, ((e1 = op(e1, e2)) | ~ sP25), inference(cnf_transformation, [], [f161])).
fof(f1212, plain, (~ spl66_104 | spl66_42), inference(avatar_split_clause, [], [f445, f764, f1206])).
fof(f445, plain, ((e1 = op(e1, e1)) | ~ sP26), inference(cnf_transformation, [], [f160])).
fof(f160, plain, (((e1 = op(e1, e3)) & (e1 = op(e3, e1)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP26), inference(nnf_transformation, [], [f80])).
fof(f1210, plain, (~ spl66_104 | spl66_10), inference(avatar_split_clause, [], [f447, f628, f1206])).
fof(f447, plain, ((e1 = op(e3, e1)) | ~ sP26), inference(cnf_transformation, [], [f160])).
fof(f1209, plain, (~ spl66_104 | spl66_34), inference(avatar_split_clause, [], [f448, f730, f1206])).
fof(f448, plain, ((e1 = op(e1, e3)) | ~ sP26), inference(cnf_transformation, [], [f160])).
fof(f1204, plain, (~ spl66_103 | spl66_38), inference(avatar_split_clause, [], [f441, f747, f1198])).
fof(f441, plain, ((e1 = op(e1, e2)) | ~ sP27), inference(cnf_transformation, [], [f159])).
fof(f159, plain, (((e2 = op(e2, e0)) & (e2 = op(e0, e2)) & (e1 = op(e2, e1)) & (e1 = op(e1, e2))) | ~ sP27), inference(nnf_transformation, [], [f81])).
fof(f1203, plain, (~ spl66_103 | spl66_26), inference(avatar_split_clause, [], [f442, f696, f1198])).
fof(f442, plain, ((e1 = op(e2, e1)) | ~ sP27), inference(cnf_transformation, [], [f159])).
fof(f1202, plain, (~ spl66_103 | spl66_55), inference(avatar_split_clause, [], [f443, f819, f1198])).
fof(f443, plain, ((e2 = op(e0, e2)) | ~ sP27), inference(cnf_transformation, [], [f159])).
fof(f1201, plain, (~ spl66_103 | spl66_31), inference(avatar_split_clause, [], [f444, f717, f1198])).
fof(f444, plain, ((e2 = op(e2, e0)) | ~ sP27), inference(cnf_transformation, [], [f159])).
fof(f1196, plain, (~ spl66_102 | spl66_38), inference(avatar_split_clause, [], [f437, f747, f1190])).
fof(f437, plain, ((e1 = op(e1, e2)) | ~ sP28), inference(cnf_transformation, [], [f158])).
fof(f158, plain, (((e2 = op(e2, e1)) & (e2 = op(e1, e2)) & (e1 = op(e2, e1)) & (e1 = op(e1, e2))) | ~ sP28), inference(nnf_transformation, [], [f82])).
fof(f1195, plain, (~ spl66_102 | spl66_26), inference(avatar_split_clause, [], [f438, f696, f1190])).
fof(f438, plain, ((e1 = op(e2, e1)) | ~ sP28), inference(cnf_transformation, [], [f158])).
fof(f1194, plain, (~ spl66_102 | spl66_39), inference(avatar_split_clause, [], [f439, f751, f1190])).
fof(f439, plain, ((e2 = op(e1, e2)) | ~ sP28), inference(cnf_transformation, [], [f158])).
fof(f1193, plain, (~ spl66_102 | spl66_27), inference(avatar_split_clause, [], [f440, f700, f1190])).
fof(f440, plain, ((e2 = op(e2, e1)) | ~ sP28), inference(cnf_transformation, [], [f158])).
fof(f1188, plain, (~ spl66_101 | spl66_38), inference(avatar_split_clause, [], [f433, f747, f1182])).
fof(f433, plain, ((e1 = op(e1, e2)) | ~ sP29), inference(cnf_transformation, [], [f157])).
fof(f157, plain, (((e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e1 = op(e2, e1)) & (e1 = op(e1, e2))) | ~ sP29), inference(nnf_transformation, [], [f83])).
fof(f1186, plain, (~ spl66_101 | spl66_23), inference(avatar_split_clause, [], [f435, f683, f1182])).
fof(f435, plain, ((e2 = op(e2, e2)) | ~ sP29), inference(cnf_transformation, [], [f157])).
fof(f1180, plain, (~ spl66_100 | spl66_38), inference(avatar_split_clause, [], [f429, f747, f1174])).
fof(f429, plain, ((e1 = op(e1, e2)) | ~ sP30), inference(cnf_transformation, [], [f156])).
fof(f156, plain, (((e2 = op(e2, e3)) & (e2 = op(e3, e2)) & (e1 = op(e2, e1)) & (e1 = op(e1, e2))) | ~ sP30), inference(nnf_transformation, [], [f84])).
fof(f1179, plain, (~ spl66_100 | spl66_26), inference(avatar_split_clause, [], [f430, f696, f1174])).
fof(f430, plain, ((e1 = op(e2, e1)) | ~ sP30), inference(cnf_transformation, [], [f156])).
fof(f1178, plain, (~ spl66_100 | spl66_7), inference(avatar_split_clause, [], [f431, f615, f1174])).
fof(f431, plain, ((e2 = op(e3, e2)) | ~ sP30), inference(cnf_transformation, [], [f156])).
fof(f1177, plain, (~ spl66_100 | spl66_19), inference(avatar_split_clause, [], [f432, f666, f1174])).
fof(f432, plain, ((e2 = op(e2, e3)) | ~ sP30), inference(cnf_transformation, [], [f156])).
fof(f1172, plain, (~ spl66_99 | spl66_34), inference(avatar_split_clause, [], [f425, f730, f1166])).
fof(f425, plain, ((e1 = op(e1, e3)) | ~ sP31), inference(cnf_transformation, [], [f155])).
fof(f155, plain, (((e3 = op(e3, e0)) & (e3 = op(e0, e3)) & (e1 = op(e3, e1)) & (e1 = op(e1, e3))) | ~ sP31), inference(nnf_transformation, [], [f85])).
fof(f1171, plain, (~ spl66_99 | spl66_10), inference(avatar_split_clause, [], [f426, f628, f1166])).
fof(f426, plain, ((e1 = op(e3, e1)) | ~ sP31), inference(cnf_transformation, [], [f155])).
fof(f1170, plain, (~ spl66_99 | spl66_52), inference(avatar_split_clause, [], [f427, f806, f1166])).
fof(f427, plain, ((e3 = op(e0, e3)) | ~ sP31), inference(cnf_transformation, [], [f155])).
fof(f1169, plain, (~ spl66_99 | spl66_16), inference(avatar_split_clause, [], [f428, f653, f1166])).
fof(f428, plain, ((e3 = op(e3, e0)) | ~ sP31), inference(cnf_transformation, [], [f155])).
fof(f1164, plain, (~ spl66_98 | spl66_34), inference(avatar_split_clause, [], [f421, f730, f1158])).
fof(f421, plain, ((e1 = op(e1, e3)) | ~ sP32), inference(cnf_transformation, [], [f154])).
fof(f154, plain, (((e3 = op(e3, e1)) & (e3 = op(e1, e3)) & (e1 = op(e3, e1)) & (e1 = op(e1, e3))) | ~ sP32), inference(nnf_transformation, [], [f86])).
fof(f1163, plain, (~ spl66_98 | spl66_10), inference(avatar_split_clause, [], [f422, f628, f1158])).
fof(f422, plain, ((e1 = op(e3, e1)) | ~ sP32), inference(cnf_transformation, [], [f154])).
fof(f1162, plain, (~ spl66_98 | spl66_36), inference(avatar_split_clause, [], [f423, f738, f1158])).
fof(f423, plain, ((e3 = op(e1, e3)) | ~ sP32), inference(cnf_transformation, [], [f154])).
fof(f1161, plain, (~ spl66_98 | spl66_12), inference(avatar_split_clause, [], [f424, f636, f1158])).
fof(f424, plain, ((e3 = op(e3, e1)) | ~ sP32), inference(cnf_transformation, [], [f154])).
fof(f1156, plain, (~ spl66_97 | spl66_34), inference(avatar_split_clause, [], [f417, f730, f1150])).
fof(f417, plain, ((e1 = op(e1, e3)) | ~ sP33), inference(cnf_transformation, [], [f153])).
fof(f153, plain, (((e3 = op(e3, e2)) & (e3 = op(e2, e3)) & (e1 = op(e3, e1)) & (e1 = op(e1, e3))) | ~ sP33), inference(nnf_transformation, [], [f87])).
fof(f1155, plain, (~ spl66_97 | spl66_10), inference(avatar_split_clause, [], [f418, f628, f1150])).
fof(f418, plain, ((e1 = op(e3, e1)) | ~ sP33), inference(cnf_transformation, [], [f153])).
fof(f1154, plain, (~ spl66_97 | spl66_20), inference(avatar_split_clause, [], [f419, f670, f1150])).
fof(f419, plain, ((e3 = op(e2, e3)) | ~ sP33), inference(cnf_transformation, [], [f153])).
fof(f1153, plain, (~ spl66_97 | spl66_8), inference(avatar_split_clause, [], [f420, f619, f1150])).
fof(f420, plain, ((e3 = op(e3, e2)) | ~ sP33), inference(cnf_transformation, [], [f153])).
fof(f1147, plain, (~ spl66_96 | spl66_10), inference(avatar_split_clause, [], [f414, f628, f1142])).
fof(f414, plain, ((e1 = op(e3, e1)) | ~ sP34), inference(cnf_transformation, [], [f152])).
fof(f152, plain, (((e3 = op(e3, e3)) & (e3 = op(e3, e3)) & (e1 = op(e3, e1)) & (e1 = op(e1, e3))) | ~ sP34), inference(nnf_transformation, [], [f88])).
fof(f1146, plain, (~ spl66_96 | spl66_4), inference(avatar_split_clause, [], [f415, f602, f1142])).
fof(f415, plain, ((e3 = op(e3, e3)) | ~ sP34), inference(cnf_transformation, [], [f152])).
fof(f1140, plain, (~ spl66_95 | spl66_31), inference(avatar_split_clause, [], [f409, f717, f1134])).
fof(f409, plain, ((e2 = op(e2, e0)) | ~ sP35), inference(cnf_transformation, [], [f151])).
fof(f151, plain, (((e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e2 = op(e0, e2)) & (e2 = op(e2, e0))) | ~ sP35), inference(nnf_transformation, [], [f89])).
fof(f1139, plain, (~ spl66_95 | spl66_55), inference(avatar_split_clause, [], [f410, f819, f1134])).
fof(f410, plain, ((e2 = op(e0, e2)) | ~ sP35), inference(cnf_transformation, [], [f151])).
fof(f1138, plain, (~ spl66_95 | spl66_61), inference(avatar_split_clause, [], [f411, f845, f1134])).
fof(f411, plain, ((e0 = op(e0, e0)) | ~ sP35), inference(cnf_transformation, [], [f151])).
fof(f1132, plain, (~ spl66_94 | spl66_31), inference(avatar_split_clause, [], [f405, f717, f1126])).
fof(f405, plain, ((e2 = op(e2, e0)) | ~ sP36), inference(cnf_transformation, [], [f150])).
fof(f150, plain, (((e0 = op(e0, e1)) & (e0 = op(e1, e0)) & (e2 = op(e0, e2)) & (e2 = op(e2, e0))) | ~ sP36), inference(nnf_transformation, [], [f90])).
fof(f1131, plain, (~ spl66_94 | spl66_55), inference(avatar_split_clause, [], [f406, f819, f1126])).
fof(f406, plain, ((e2 = op(e0, e2)) | ~ sP36), inference(cnf_transformation, [], [f150])).
fof(f1130, plain, (~ spl66_94 | spl66_45), inference(avatar_split_clause, [], [f407, f777, f1126])).
fof(f407, plain, ((e0 = op(e1, e0)) | ~ sP36), inference(cnf_transformation, [], [f150])).
fof(f1129, plain, (~ spl66_94 | spl66_57), inference(avatar_split_clause, [], [f408, f828, f1126])).
fof(f408, plain, ((e0 = op(e0, e1)) | ~ sP36), inference(cnf_transformation, [], [f150])).
fof(f1124, plain, (~ spl66_93 | spl66_31), inference(avatar_split_clause, [], [f401, f717, f1118])).
fof(f401, plain, ((e2 = op(e2, e0)) | ~ sP37), inference(cnf_transformation, [], [f149])).
fof(f149, plain, (((e0 = op(e0, e2)) & (e0 = op(e2, e0)) & (e2 = op(e0, e2)) & (e2 = op(e2, e0))) | ~ sP37), inference(nnf_transformation, [], [f91])).
fof(f1123, plain, (~ spl66_93 | spl66_55), inference(avatar_split_clause, [], [f402, f819, f1118])).
fof(f402, plain, ((e2 = op(e0, e2)) | ~ sP37), inference(cnf_transformation, [], [f149])).
fof(f1121, plain, (~ spl66_93 | spl66_53), inference(avatar_split_clause, [], [f404, f811, f1118])).
fof(f404, plain, ((e0 = op(e0, e2)) | ~ sP37), inference(cnf_transformation, [], [f149])).
fof(f1116, plain, (~ spl66_92 | spl66_31), inference(avatar_split_clause, [], [f397, f717, f1110])).
fof(f397, plain, ((e2 = op(e2, e0)) | ~ sP38), inference(cnf_transformation, [], [f148])).
fof(f148, plain, (((e0 = op(e0, e3)) & (e0 = op(e3, e0)) & (e2 = op(e0, e2)) & (e2 = op(e2, e0))) | ~ sP38), inference(nnf_transformation, [], [f92])).
fof(f1115, plain, (~ spl66_92 | spl66_55), inference(avatar_split_clause, [], [f398, f819, f1110])).
fof(f398, plain, ((e2 = op(e0, e2)) | ~ sP38), inference(cnf_transformation, [], [f148])).
fof(f1114, plain, (~ spl66_92 | spl66_13), inference(avatar_split_clause, [], [f399, f641, f1110])).
fof(f399, plain, ((e0 = op(e3, e0)) | ~ sP38), inference(cnf_transformation, [], [f148])).
fof(f1113, plain, (~ spl66_92 | spl66_49), inference(avatar_split_clause, [], [f400, f794, f1110])).
fof(f400, plain, ((e0 = op(e0, e3)) | ~ sP38), inference(cnf_transformation, [], [f148])).
fof(f1108, plain, (~ spl66_91 | spl66_27), inference(avatar_split_clause, [], [f393, f700, f1102])).
fof(f393, plain, ((e2 = op(e2, e1)) | ~ sP39), inference(cnf_transformation, [], [f147])).
fof(f147, plain, (((e1 = op(e1, e0)) & (e1 = op(e0, e1)) & (e2 = op(e1, e2)) & (e2 = op(e2, e1))) | ~ sP39), inference(nnf_transformation, [], [f93])).
fof(f1107, plain, (~ spl66_91 | spl66_39), inference(avatar_split_clause, [], [f394, f751, f1102])).
fof(f394, plain, ((e2 = op(e1, e2)) | ~ sP39), inference(cnf_transformation, [], [f147])).
fof(f1106, plain, (~ spl66_91 | spl66_58), inference(avatar_split_clause, [], [f395, f832, f1102])).
fof(f395, plain, ((e1 = op(e0, e1)) | ~ sP39), inference(cnf_transformation, [], [f147])).
fof(f1105, plain, (~ spl66_91 | spl66_46), inference(avatar_split_clause, [], [f396, f781, f1102])).
fof(f396, plain, ((e1 = op(e1, e0)) | ~ sP39), inference(cnf_transformation, [], [f147])).
fof(f1100, plain, (~ spl66_90 | spl66_27), inference(avatar_split_clause, [], [f389, f700, f1094])).
fof(f389, plain, ((e2 = op(e2, e1)) | ~ sP40), inference(cnf_transformation, [], [f146])).
fof(f146, plain, (((e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e2 = op(e1, e2)) & (e2 = op(e2, e1))) | ~ sP40), inference(nnf_transformation, [], [f94])).
fof(f1099, plain, (~ spl66_90 | spl66_39), inference(avatar_split_clause, [], [f390, f751, f1094])).
fof(f390, plain, ((e2 = op(e1, e2)) | ~ sP40), inference(cnf_transformation, [], [f146])).
fof(f1098, plain, (~ spl66_90 | spl66_42), inference(avatar_split_clause, [], [f391, f764, f1094])).
fof(f391, plain, ((e1 = op(e1, e1)) | ~ sP40), inference(cnf_transformation, [], [f146])).
fof(f1092, plain, (~ spl66_89 | spl66_27), inference(avatar_split_clause, [], [f385, f700, f1086])).
fof(f385, plain, ((e2 = op(e2, e1)) | ~ sP41), inference(cnf_transformation, [], [f145])).
fof(f145, plain, (((e1 = op(e1, e2)) & (e1 = op(e2, e1)) & (e2 = op(e1, e2)) & (e2 = op(e2, e1))) | ~ sP41), inference(nnf_transformation, [], [f95])).
fof(f1091, plain, (~ spl66_89 | spl66_39), inference(avatar_split_clause, [], [f386, f751, f1086])).
fof(f386, plain, ((e2 = op(e1, e2)) | ~ sP41), inference(cnf_transformation, [], [f145])).
fof(f1090, plain, (~ spl66_89 | spl66_26), inference(avatar_split_clause, [], [f387, f696, f1086])).
fof(f387, plain, ((e1 = op(e2, e1)) | ~ sP41), inference(cnf_transformation, [], [f145])).
fof(f1089, plain, (~ spl66_89 | spl66_38), inference(avatar_split_clause, [], [f388, f747, f1086])).
fof(f388, plain, ((e1 = op(e1, e2)) | ~ sP41), inference(cnf_transformation, [], [f145])).
fof(f1084, plain, (~ spl66_88 | spl66_27), inference(avatar_split_clause, [], [f381, f700, f1078])).
fof(f381, plain, ((e2 = op(e2, e1)) | ~ sP42), inference(cnf_transformation, [], [f144])).
fof(f144, plain, (((e1 = op(e1, e3)) & (e1 = op(e3, e1)) & (e2 = op(e1, e2)) & (e2 = op(e2, e1))) | ~ sP42), inference(nnf_transformation, [], [f96])).
fof(f1083, plain, (~ spl66_88 | spl66_39), inference(avatar_split_clause, [], [f382, f751, f1078])).
fof(f382, plain, ((e2 = op(e1, e2)) | ~ sP42), inference(cnf_transformation, [], [f144])).
fof(f1082, plain, (~ spl66_88 | spl66_10), inference(avatar_split_clause, [], [f383, f628, f1078])).
fof(f383, plain, ((e1 = op(e3, e1)) | ~ sP42), inference(cnf_transformation, [], [f144])).
fof(f1081, plain, (~ spl66_88 | spl66_34), inference(avatar_split_clause, [], [f384, f730, f1078])).
fof(f384, plain, ((e1 = op(e1, e3)) | ~ sP42), inference(cnf_transformation, [], [f144])).
fof(f1076, plain, (~ spl66_87 | spl66_23), inference(avatar_split_clause, [], [f377, f683, f1070])).
fof(f377, plain, ((e2 = op(e2, e2)) | ~ sP43), inference(cnf_transformation, [], [f143])).
fof(f143, plain, (((e2 = op(e2, e0)) & (e2 = op(e0, e2)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP43), inference(nnf_transformation, [], [f97])).
fof(f1074, plain, (~ spl66_87 | spl66_55), inference(avatar_split_clause, [], [f379, f819, f1070])).
fof(f379, plain, ((e2 = op(e0, e2)) | ~ sP43), inference(cnf_transformation, [], [f143])).
fof(f1068, plain, (~ spl66_86 | spl66_23), inference(avatar_split_clause, [], [f373, f683, f1062])).
fof(f373, plain, ((e2 = op(e2, e2)) | ~ sP44), inference(cnf_transformation, [], [f142])).
fof(f142, plain, (((e2 = op(e2, e1)) & (e2 = op(e1, e2)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP44), inference(nnf_transformation, [], [f98])).
fof(f1060, plain, (~ spl66_85 | spl66_23), inference(avatar_split_clause, [], [f369, f683, f1054])).
fof(f369, plain, ((e2 = op(e2, e2)) | ~ sP45), inference(cnf_transformation, [], [f141])).
fof(f141, plain, (((e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP45), inference(nnf_transformation, [], [f99])).
fof(f1052, plain, (~ spl66_84 | spl66_23), inference(avatar_split_clause, [], [f365, f683, f1046])).
fof(f365, plain, ((e2 = op(e2, e2)) | ~ sP46), inference(cnf_transformation, [], [f140])).
fof(f140, plain, (((e2 = op(e2, e3)) & (e2 = op(e3, e2)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP46), inference(nnf_transformation, [], [f100])).
fof(f1050, plain, (~ spl66_84 | spl66_7), inference(avatar_split_clause, [], [f367, f615, f1046])).
fof(f367, plain, ((e2 = op(e3, e2)) | ~ sP46), inference(cnf_transformation, [], [f140])).
fof(f1049, plain, (~ spl66_84 | spl66_19), inference(avatar_split_clause, [], [f368, f666, f1046])).
fof(f368, plain, ((e2 = op(e2, e3)) | ~ sP46), inference(cnf_transformation, [], [f140])).
fof(f1044, plain, (~ spl66_83 | spl66_19), inference(avatar_split_clause, [], [f361, f666, f1038])).
fof(f361, plain, ((e2 = op(e2, e3)) | ~ sP47), inference(cnf_transformation, [], [f139])).
fof(f139, plain, (((e3 = op(e3, e0)) & (e3 = op(e0, e3)) & (e2 = op(e3, e2)) & (e2 = op(e2, e3))) | ~ sP47), inference(nnf_transformation, [], [f101])).
fof(f1043, plain, (~ spl66_83 | spl66_7), inference(avatar_split_clause, [], [f362, f615, f1038])).
fof(f362, plain, ((e2 = op(e3, e2)) | ~ sP47), inference(cnf_transformation, [], [f139])).
fof(f1042, plain, (~ spl66_83 | spl66_52), inference(avatar_split_clause, [], [f363, f806, f1038])).
fof(f363, plain, ((e3 = op(e0, e3)) | ~ sP47), inference(cnf_transformation, [], [f139])).
fof(f1041, plain, (~ spl66_83 | spl66_16), inference(avatar_split_clause, [], [f364, f653, f1038])).
fof(f364, plain, ((e3 = op(e3, e0)) | ~ sP47), inference(cnf_transformation, [], [f139])).
fof(f1036, plain, (~ spl66_82 | spl66_19), inference(avatar_split_clause, [], [f357, f666, f1030])).
fof(f357, plain, ((e2 = op(e2, e3)) | ~ sP48), inference(cnf_transformation, [], [f138])).
fof(f138, plain, (((e3 = op(e3, e1)) & (e3 = op(e1, e3)) & (e2 = op(e3, e2)) & (e2 = op(e2, e3))) | ~ sP48), inference(nnf_transformation, [], [f102])).
fof(f1035, plain, (~ spl66_82 | spl66_7), inference(avatar_split_clause, [], [f358, f615, f1030])).
fof(f358, plain, ((e2 = op(e3, e2)) | ~ sP48), inference(cnf_transformation, [], [f138])).
fof(f1034, plain, (~ spl66_82 | spl66_36), inference(avatar_split_clause, [], [f359, f738, f1030])).
fof(f359, plain, ((e3 = op(e1, e3)) | ~ sP48), inference(cnf_transformation, [], [f138])).
fof(f1033, plain, (~ spl66_82 | spl66_12), inference(avatar_split_clause, [], [f360, f636, f1030])).
fof(f360, plain, ((e3 = op(e3, e1)) | ~ sP48), inference(cnf_transformation, [], [f138])).
fof(f1028, plain, (~ spl66_81 | spl66_19), inference(avatar_split_clause, [], [f353, f666, f1022])).
fof(f353, plain, ((e2 = op(e2, e3)) | ~ sP49), inference(cnf_transformation, [], [f137])).
fof(f137, plain, (((e3 = op(e3, e2)) & (e3 = op(e2, e3)) & (e2 = op(e3, e2)) & (e2 = op(e2, e3))) | ~ sP49), inference(nnf_transformation, [], [f103])).
fof(f1027, plain, (~ spl66_81 | spl66_7), inference(avatar_split_clause, [], [f354, f615, f1022])).
fof(f354, plain, ((e2 = op(e3, e2)) | ~ sP49), inference(cnf_transformation, [], [f137])).
fof(f1026, plain, (~ spl66_81 | spl66_20), inference(avatar_split_clause, [], [f355, f670, f1022])).
fof(f355, plain, ((e3 = op(e2, e3)) | ~ sP49), inference(cnf_transformation, [], [f137])).
fof(f1025, plain, (~ spl66_81 | spl66_8), inference(avatar_split_clause, [], [f356, f619, f1022])).
fof(f356, plain, ((e3 = op(e3, e2)) | ~ sP49), inference(cnf_transformation, [], [f137])).
fof(f1020, plain, (~ spl66_80 | spl66_19), inference(avatar_split_clause, [], [f349, f666, f1014])).
fof(f349, plain, ((e2 = op(e2, e3)) | ~ sP50), inference(cnf_transformation, [], [f136])).
fof(f136, plain, (((e3 = op(e3, e3)) & (e3 = op(e3, e3)) & (e2 = op(e3, e2)) & (e2 = op(e2, e3))) | ~ sP50), inference(nnf_transformation, [], [f104])).
fof(f1019, plain, (~ spl66_80 | spl66_7), inference(avatar_split_clause, [], [f350, f615, f1014])).
fof(f350, plain, ((e2 = op(e3, e2)) | ~ sP50), inference(cnf_transformation, [], [f136])).
fof(f1018, plain, (~ spl66_80 | spl66_4), inference(avatar_split_clause, [], [f351, f602, f1014])).
fof(f351, plain, ((e3 = op(e3, e3)) | ~ sP50), inference(cnf_transformation, [], [f136])).
fof(f1012, plain, (~ spl66_79 | spl66_16), inference(avatar_split_clause, [], [f345, f653, f1006])).
fof(f345, plain, ((e3 = op(e3, e0)) | ~ sP51), inference(cnf_transformation, [], [f135])).
fof(f135, plain, (((e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e3 = op(e0, e3)) & (e3 = op(e3, e0))) | ~ sP51), inference(nnf_transformation, [], [f105])).
fof(f1011, plain, (~ spl66_79 | spl66_52), inference(avatar_split_clause, [], [f346, f806, f1006])).
fof(f346, plain, ((e3 = op(e0, e3)) | ~ sP51), inference(cnf_transformation, [], [f135])).
fof(f1010, plain, (~ spl66_79 | spl66_61), inference(avatar_split_clause, [], [f347, f845, f1006])).
fof(f347, plain, ((e0 = op(e0, e0)) | ~ sP51), inference(cnf_transformation, [], [f135])).
fof(f1004, plain, (~ spl66_78 | spl66_16), inference(avatar_split_clause, [], [f341, f653, f998])).
fof(f341, plain, ((e3 = op(e3, e0)) | ~ sP52), inference(cnf_transformation, [], [f134])).
fof(f134, plain, (((e0 = op(e0, e1)) & (e0 = op(e1, e0)) & (e3 = op(e0, e3)) & (e3 = op(e3, e0))) | ~ sP52), inference(nnf_transformation, [], [f106])).
fof(f1003, plain, (~ spl66_78 | spl66_52), inference(avatar_split_clause, [], [f342, f806, f998])).
fof(f342, plain, ((e3 = op(e0, e3)) | ~ sP52), inference(cnf_transformation, [], [f134])).
fof(f1002, plain, (~ spl66_78 | spl66_45), inference(avatar_split_clause, [], [f343, f777, f998])).
fof(f343, plain, ((e0 = op(e1, e0)) | ~ sP52), inference(cnf_transformation, [], [f134])).
fof(f1001, plain, (~ spl66_78 | spl66_57), inference(avatar_split_clause, [], [f344, f828, f998])).
fof(f344, plain, ((e0 = op(e0, e1)) | ~ sP52), inference(cnf_transformation, [], [f134])).
fof(f996, plain, (~ spl66_77 | spl66_16), inference(avatar_split_clause, [], [f337, f653, f990])).
fof(f337, plain, ((e3 = op(e3, e0)) | ~ sP53), inference(cnf_transformation, [], [f133])).
fof(f133, plain, (((e0 = op(e0, e2)) & (e0 = op(e2, e0)) & (e3 = op(e0, e3)) & (e3 = op(e3, e0))) | ~ sP53), inference(nnf_transformation, [], [f107])).
fof(f995, plain, (~ spl66_77 | spl66_52), inference(avatar_split_clause, [], [f338, f806, f990])).
fof(f338, plain, ((e3 = op(e0, e3)) | ~ sP53), inference(cnf_transformation, [], [f133])).
fof(f994, plain, (~ spl66_77 | spl66_29), inference(avatar_split_clause, [], [f339, f709, f990])).
fof(f339, plain, ((e0 = op(e2, e0)) | ~ sP53), inference(cnf_transformation, [], [f133])).
fof(f993, plain, (~ spl66_77 | spl66_53), inference(avatar_split_clause, [], [f340, f811, f990])).
fof(f340, plain, ((e0 = op(e0, e2)) | ~ sP53), inference(cnf_transformation, [], [f133])).
fof(f988, plain, (~ spl66_76 | spl66_16), inference(avatar_split_clause, [], [f333, f653, f982])).
fof(f333, plain, ((e3 = op(e3, e0)) | ~ sP54), inference(cnf_transformation, [], [f132])).
fof(f132, plain, (((e0 = op(e0, e3)) & (e0 = op(e3, e0)) & (e3 = op(e0, e3)) & (e3 = op(e3, e0))) | ~ sP54), inference(nnf_transformation, [], [f108])).
fof(f987, plain, (~ spl66_76 | spl66_52), inference(avatar_split_clause, [], [f334, f806, f982])).
fof(f334, plain, ((e3 = op(e0, e3)) | ~ sP54), inference(cnf_transformation, [], [f132])).
fof(f986, plain, (~ spl66_76 | spl66_13), inference(avatar_split_clause, [], [f335, f641, f982])).
fof(f335, plain, ((e0 = op(e3, e0)) | ~ sP54), inference(cnf_transformation, [], [f132])).
fof(f985, plain, (~ spl66_76 | spl66_49), inference(avatar_split_clause, [], [f336, f794, f982])).
fof(f336, plain, ((e0 = op(e0, e3)) | ~ sP54), inference(cnf_transformation, [], [f132])).
fof(f980, plain, (~ spl66_75 | spl66_12), inference(avatar_split_clause, [], [f329, f636, f974])).
fof(f329, plain, ((e3 = op(e3, e1)) | ~ sP55), inference(cnf_transformation, [], [f131])).
fof(f131, plain, (((e1 = op(e1, e0)) & (e1 = op(e0, e1)) & (e3 = op(e1, e3)) & (e3 = op(e3, e1))) | ~ sP55), inference(nnf_transformation, [], [f109])).
fof(f979, plain, (~ spl66_75 | spl66_36), inference(avatar_split_clause, [], [f330, f738, f974])).
fof(f330, plain, ((e3 = op(e1, e3)) | ~ sP55), inference(cnf_transformation, [], [f131])).
fof(f978, plain, (~ spl66_75 | spl66_58), inference(avatar_split_clause, [], [f331, f832, f974])).
fof(f331, plain, ((e1 = op(e0, e1)) | ~ sP55), inference(cnf_transformation, [], [f131])).
fof(f977, plain, (~ spl66_75 | spl66_46), inference(avatar_split_clause, [], [f332, f781, f974])).
fof(f332, plain, ((e1 = op(e1, e0)) | ~ sP55), inference(cnf_transformation, [], [f131])).
fof(f972, plain, (~ spl66_74 | spl66_12), inference(avatar_split_clause, [], [f325, f636, f966])).
fof(f325, plain, ((e3 = op(e3, e1)) | ~ sP56), inference(cnf_transformation, [], [f130])).
fof(f130, plain, (((e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e3 = op(e1, e3)) & (e3 = op(e3, e1))) | ~ sP56), inference(nnf_transformation, [], [f110])).
fof(f971, plain, (~ spl66_74 | spl66_36), inference(avatar_split_clause, [], [f326, f738, f966])).
fof(f326, plain, ((e3 = op(e1, e3)) | ~ sP56), inference(cnf_transformation, [], [f130])).
fof(f970, plain, (~ spl66_74 | spl66_42), inference(avatar_split_clause, [], [f327, f764, f966])).
fof(f327, plain, ((e1 = op(e1, e1)) | ~ sP56), inference(cnf_transformation, [], [f130])).
fof(f964, plain, (~ spl66_73 | spl66_12), inference(avatar_split_clause, [], [f321, f636, f958])).
fof(f321, plain, ((e3 = op(e3, e1)) | ~ sP57), inference(cnf_transformation, [], [f129])).
fof(f129, plain, (((e1 = op(e1, e2)) & (e1 = op(e2, e1)) & (e3 = op(e1, e3)) & (e3 = op(e3, e1))) | ~ sP57), inference(nnf_transformation, [], [f111])).
fof(f963, plain, (~ spl66_73 | spl66_36), inference(avatar_split_clause, [], [f322, f738, f958])).
fof(f322, plain, ((e3 = op(e1, e3)) | ~ sP57), inference(cnf_transformation, [], [f129])).
fof(f962, plain, (~ spl66_73 | spl66_26), inference(avatar_split_clause, [], [f323, f696, f958])).
fof(f323, plain, ((e1 = op(e2, e1)) | ~ sP57), inference(cnf_transformation, [], [f129])).
fof(f961, plain, (~ spl66_73 | spl66_38), inference(avatar_split_clause, [], [f324, f747, f958])).
fof(f324, plain, ((e1 = op(e1, e2)) | ~ sP57), inference(cnf_transformation, [], [f129])).
fof(f956, plain, (~ spl66_72 | spl66_12), inference(avatar_split_clause, [], [f317, f636, f950])).
fof(f317, plain, ((e3 = op(e3, e1)) | ~ sP58), inference(cnf_transformation, [], [f128])).
fof(f128, plain, (((e1 = op(e1, e3)) & (e1 = op(e3, e1)) & (e3 = op(e1, e3)) & (e3 = op(e3, e1))) | ~ sP58), inference(nnf_transformation, [], [f112])).
fof(f955, plain, (~ spl66_72 | spl66_36), inference(avatar_split_clause, [], [f318, f738, f950])).
fof(f318, plain, ((e3 = op(e1, e3)) | ~ sP58), inference(cnf_transformation, [], [f128])).
fof(f954, plain, (~ spl66_72 | spl66_10), inference(avatar_split_clause, [], [f319, f628, f950])).
fof(f319, plain, ((e1 = op(e3, e1)) | ~ sP58), inference(cnf_transformation, [], [f128])).
fof(f953, plain, (~ spl66_72 | spl66_34), inference(avatar_split_clause, [], [f320, f730, f950])).
fof(f320, plain, ((e1 = op(e1, e3)) | ~ sP58), inference(cnf_transformation, [], [f128])).
fof(f948, plain, (~ spl66_71 | spl66_8), inference(avatar_split_clause, [], [f313, f619, f942])).
fof(f313, plain, ((e3 = op(e3, e2)) | ~ sP59), inference(cnf_transformation, [], [f127])).
fof(f127, plain, (((e2 = op(e2, e0)) & (e2 = op(e0, e2)) & (e3 = op(e2, e3)) & (e3 = op(e3, e2))) | ~ sP59), inference(nnf_transformation, [], [f113])).
fof(f947, plain, (~ spl66_71 | spl66_20), inference(avatar_split_clause, [], [f314, f670, f942])).
fof(f314, plain, ((e3 = op(e2, e3)) | ~ sP59), inference(cnf_transformation, [], [f127])).
fof(f946, plain, (~ spl66_71 | spl66_55), inference(avatar_split_clause, [], [f315, f819, f942])).
fof(f315, plain, ((e2 = op(e0, e2)) | ~ sP59), inference(cnf_transformation, [], [f127])).
fof(f945, plain, (~ spl66_71 | spl66_31), inference(avatar_split_clause, [], [f316, f717, f942])).
fof(f316, plain, ((e2 = op(e2, e0)) | ~ sP59), inference(cnf_transformation, [], [f127])).
fof(f940, plain, (~ spl66_70 | spl66_8), inference(avatar_split_clause, [], [f309, f619, f934])).
fof(f309, plain, ((e3 = op(e3, e2)) | ~ sP60), inference(cnf_transformation, [], [f126])).
fof(f126, plain, (((e2 = op(e2, e1)) & (e2 = op(e1, e2)) & (e3 = op(e2, e3)) & (e3 = op(e3, e2))) | ~ sP60), inference(nnf_transformation, [], [f114])).
fof(f939, plain, (~ spl66_70 | spl66_20), inference(avatar_split_clause, [], [f310, f670, f934])).
fof(f310, plain, ((e3 = op(e2, e3)) | ~ sP60), inference(cnf_transformation, [], [f126])).
fof(f938, plain, (~ spl66_70 | spl66_39), inference(avatar_split_clause, [], [f311, f751, f934])).
fof(f311, plain, ((e2 = op(e1, e2)) | ~ sP60), inference(cnf_transformation, [], [f126])).
fof(f937, plain, (~ spl66_70 | spl66_27), inference(avatar_split_clause, [], [f312, f700, f934])).
fof(f312, plain, ((e2 = op(e2, e1)) | ~ sP60), inference(cnf_transformation, [], [f126])).
fof(f932, plain, (~ spl66_69 | spl66_8), inference(avatar_split_clause, [], [f305, f619, f926])).
fof(f305, plain, ((e3 = op(e3, e2)) | ~ sP61), inference(cnf_transformation, [], [f125])).
fof(f125, plain, (((e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e3 = op(e2, e3)) & (e3 = op(e3, e2))) | ~ sP61), inference(nnf_transformation, [], [f115])).
fof(f931, plain, (~ spl66_69 | spl66_20), inference(avatar_split_clause, [], [f306, f670, f926])).
fof(f306, plain, ((e3 = op(e2, e3)) | ~ sP61), inference(cnf_transformation, [], [f125])).
fof(f930, plain, (~ spl66_69 | spl66_23), inference(avatar_split_clause, [], [f307, f683, f926])).
fof(f307, plain, ((e2 = op(e2, e2)) | ~ sP61), inference(cnf_transformation, [], [f125])).
fof(f924, plain, (~ spl66_68 | spl66_8), inference(avatar_split_clause, [], [f301, f619, f918])).
fof(f301, plain, ((e3 = op(e3, e2)) | ~ sP62), inference(cnf_transformation, [], [f124])).
fof(f124, plain, (((e2 = op(e2, e3)) & (e2 = op(e3, e2)) & (e3 = op(e2, e3)) & (e3 = op(e3, e2))) | ~ sP62), inference(nnf_transformation, [], [f116])).
fof(f923, plain, (~ spl66_68 | spl66_20), inference(avatar_split_clause, [], [f302, f670, f918])).
fof(f302, plain, ((e3 = op(e2, e3)) | ~ sP62), inference(cnf_transformation, [], [f124])).
fof(f922, plain, (~ spl66_68 | spl66_7), inference(avatar_split_clause, [], [f303, f615, f918])).
fof(f303, plain, ((e2 = op(e3, e2)) | ~ sP62), inference(cnf_transformation, [], [f124])).
fof(f921, plain, (~ spl66_68 | spl66_19), inference(avatar_split_clause, [], [f304, f666, f918])).
fof(f304, plain, ((e2 = op(e2, e3)) | ~ sP62), inference(cnf_transformation, [], [f124])).
fof(f916, plain, (~ spl66_67 | spl66_4), inference(avatar_split_clause, [], [f297, f602, f910])).
fof(f297, plain, ((e3 = op(e3, e3)) | ~ sP63), inference(cnf_transformation, [], [f123])).
fof(f123, plain, (((e3 = op(e3, e0)) & (e3 = op(e0, e3)) & (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ~ sP63), inference(nnf_transformation, [], [f117])).
fof(f913, plain, (~ spl66_67 | spl66_16), inference(avatar_split_clause, [], [f300, f653, f910])).
fof(f300, plain, ((e3 = op(e3, e0)) | ~ sP63), inference(cnf_transformation, [], [f123])).
fof(f908, plain, (~ spl66_66 | spl66_4), inference(avatar_split_clause, [], [f293, f602, f902])).
fof(f293, plain, ((e3 = op(e3, e3)) | ~ sP64), inference(cnf_transformation, [], [f122])).
fof(f122, plain, (((e3 = op(e3, e1)) & (e3 = op(e1, e3)) & (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ~ sP64), inference(nnf_transformation, [], [f118])).
fof(f906, plain, (~ spl66_66 | spl66_36), inference(avatar_split_clause, [], [f295, f738, f902])).
fof(f295, plain, ((e3 = op(e1, e3)) | ~ sP64), inference(cnf_transformation, [], [f122])).
fof(f905, plain, (~ spl66_66 | spl66_12), inference(avatar_split_clause, [], [f296, f636, f902])).
fof(f296, plain, ((e3 = op(e3, e1)) | ~ sP64), inference(cnf_transformation, [], [f122])).
fof(f900, plain, (~ spl66_65 | spl66_4), inference(avatar_split_clause, [], [f289, f602, f894])).
fof(f289, plain, ((e3 = op(e3, e3)) | ~ sP65), inference(cnf_transformation, [], [f121])).
fof(f121, plain, (((e3 = op(e3, e2)) & (e3 = op(e2, e3)) & (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ~ sP65), inference(nnf_transformation, [], [f119])).
fof(f898, plain, (~ spl66_65 | spl66_20), inference(avatar_split_clause, [], [f291, f670, f894])).
fof(f291, plain, ((e3 = op(e2, e3)) | ~ sP65), inference(cnf_transformation, [], [f121])).
fof(f897, plain, (~ spl66_65 | spl66_8), inference(avatar_split_clause, [], [f292, f619, f894])).
fof(f292, plain, ((e3 = op(e3, e2)) | ~ sP65), inference(cnf_transformation, [], [f121])).
fof(f891, plain, (spl66_61 | spl66_45 | spl66_29 | spl66_13), inference(avatar_split_clause, [], [f204, f641, f709, f777, f845])).
fof(f204, plain, ((e0 = op(e3, e0)) | (e0 = op(e2, e0)) | (e0 = op(e1, e0)) | (e0 = op(e0, e0))), inference(cnf_transformation, [], [f2])).
fof(f2, plain, (((e3 = op(e3, e3)) | (e3 = op(e2, e3)) | (e3 = op(e1, e3)) | (e3 = op(e0, e3))) & ((e3 = op(e3, e3)) | (e3 = op(e3, e2)) | (e3 = op(e3, e1)) | (e3 = op(e3, e0))) & ((e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(e0, e3))) & ((e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, e0))) & ((e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(e0, e3))) & ((e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, e0))) & ((e0 = op(e3, e3)) | (e0 = op(e2, e3)) | (e0 = op(e1, e3)) | (e0 = op(e0, e3))) & ((e0 = op(e3, e3)) | (e0 = op(e3, e2)) | (e0 = op(e3, e1)) | (e0 = op(e3, e0))) & ((e3 = op(e3, e2)) | (e3 = op(e2, e2)) | (e3 = op(e1, e2)) | (e3 = op(e0, e2))) & ((e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, e0))) & ((e2 = op(e3, e2)) | (e2 = op(e2, e2)) | (e2 = op(e1, e2)) | (e2 = op(e0, e2))) & ((e2 = op(e2, e3)) | (e2 = op(e2, e2)) | (e2 = op(e2, e1)) | (e2 = op(e2, e0))) & ((e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(e0, e2))) & ((e1 = op(e2, e3)) | (e1 = op(e2, e2)) | (e1 = op(e2, e1)) | (e1 = op(e2, e0))) & ((e0 = op(e3, e2)) | (e0 = op(e2, e2)) | (e0 = op(e1, e2)) | (e0 = op(e0, e2))) & ((e0 = op(e2, e3)) | (e0 = op(e2, e2)) | (e0 = op(e2, e1)) | (e0 = op(e2, e0))) & ((e3 = op(e3, e1)) | (e3 = op(e2, e1)) | (e3 = op(e1, e1)) | (e3 = op(e0, e1))) & ((e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, e0))) & ((e2 = op(e3, e1)) | (e2 = op(e2, e1)) | (e2 = op(e1, e1)) | (e2 = op(e0, e1))) & ((e2 = op(e1, e3)) | (e2 = op(e1, e2)) | (e2 = op(e1, e1)) | (e2 = op(e1, e0))) & ((e1 = op(e3, e1)) | (e1 = op(e2, e1)) | (e1 = op(e1, e1)) | (e1 = op(e0, e1))) & ((e1 = op(e1, e3)) | (e1 = op(e1, e2)) | (e1 = op(e1, e1)) | (e1 = op(e1, e0))) & ((e0 = op(e3, e1)) | (e0 = op(e2, e1)) | (e0 = op(e1, e1)) | (e0 = op(e0, e1))) & ((e0 = op(e1, e3)) | (e0 = op(e1, e2)) | (e0 = op(e1, e1)) | (e0 = op(e1, e0))) & ((e3 = op(e3, e0)) | (e3 = op(e2, e0)) | (e3 = op(e1, e0)) | (op(e0, e0) = e3)) & ((e3 = op(e0, e3)) | (e3 = op(e0, e2)) | (e3 = op(e0, e1)) | (op(e0, e0) = e3)) & ((e2 = op(e3, e0)) | (e2 = op(e2, e0)) | (e2 = op(e1, e0)) | (op(e0, e0) = e2)) & ((e2 = op(e0, e3)) | (e2 = op(e0, e2)) | (e2 = op(e0, e1)) | (op(e0, e0) = e2)) & ((e1 = op(e3, e0)) | (e1 = op(e2, e0)) | (e1 = op(e1, e0)) | (op(e0, e0) = e1)) & ((e1 = op(e0, e3)) | (e1 = op(e0, e2)) | (e1 = op(e0, e1)) | (op(e0, e0) = e1)) & ((e0 = op(e3, e0)) | (e0 = op(e2, e0)) | (e0 = op(e1, e0)) | (e0 = op(e0, e0))) & ((e0 = op(e0, e3)) | (e0 = op(e0, e2)) | (e0 = op(e0, e1)) | (e0 = op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG140+1.p', ax2)).
fof(f890, plain, (spl66_62 | spl66_58 | spl66_54 | spl66_50), inference(avatar_split_clause, [], [f205, f798, f815, f832, f849])).
fof(f205, plain, ((e1 = op(e0, e3)) | (e1 = op(e0, e2)) | (e1 = op(e0, e1)) | (op(e0, e0) = e1)), inference(cnf_transformation, [], [f2])).
fof(f889, plain, (spl66_62 | spl66_46 | spl66_30 | spl66_14), inference(avatar_split_clause, [], [f206, f645, f713, f781, f849])).
fof(f206, plain, ((e1 = op(e3, e0)) | (e1 = op(e2, e0)) | (e1 = op(e1, e0)) | (op(e0, e0) = e1)), inference(cnf_transformation, [], [f2])).
fof(f888, plain, (spl66_63 | spl66_59 | spl66_55 | spl66_51), inference(avatar_split_clause, [], [f207, f802, f819, f836, f853])).
fof(f207, plain, ((e2 = op(e0, e3)) | (e2 = op(e0, e2)) | (e2 = op(e0, e1)) | (op(e0, e0) = e2)), inference(cnf_transformation, [], [f2])).
fof(f887, plain, (spl66_63 | spl66_47 | spl66_31 | spl66_15), inference(avatar_split_clause, [], [f208, f649, f717, f785, f853])).
fof(f208, plain, ((e2 = op(e3, e0)) | (e2 = op(e2, e0)) | (e2 = op(e1, e0)) | (op(e0, e0) = e2)), inference(cnf_transformation, [], [f2])).
fof(f886, plain, (spl66_64 | spl66_60 | spl66_56 | spl66_52), inference(avatar_split_clause, [], [f209, f806, f823, f840, f857])).
fof(f209, plain, ((e3 = op(e0, e3)) | (e3 = op(e0, e2)) | (e3 = op(e0, e1)) | (op(e0, e0) = e3)), inference(cnf_transformation, [], [f2])).
fof(f885, plain, (spl66_64 | spl66_48 | spl66_32 | spl66_16), inference(avatar_split_clause, [], [f210, f653, f721, f789, f857])).
fof(f210, plain, ((e3 = op(e3, e0)) | (e3 = op(e2, e0)) | (e3 = op(e1, e0)) | (op(e0, e0) = e3)), inference(cnf_transformation, [], [f2])).
fof(f883, plain, (spl66_57 | spl66_41 | spl66_25 | spl66_9), inference(avatar_split_clause, [], [f212, f624, f692, f760, f828])).
fof(f212, plain, ((e0 = op(e3, e1)) | (e0 = op(e2, e1)) | (e0 = op(e1, e1)) | (e0 = op(e0, e1))), inference(cnf_transformation, [], [f2])).
fof(f882, plain, (spl66_46 | spl66_42 | spl66_38 | spl66_34), inference(avatar_split_clause, [], [f213, f730, f747, f764, f781])).
fof(f213, plain, ((e1 = op(e1, e3)) | (e1 = op(e1, e2)) | (e1 = op(e1, e1)) | (e1 = op(e1, e0))), inference(cnf_transformation, [], [f2])).
fof(f881, plain, (spl66_58 | spl66_42 | spl66_26 | spl66_10), inference(avatar_split_clause, [], [f214, f628, f696, f764, f832])).
fof(f214, plain, ((e1 = op(e3, e1)) | (e1 = op(e2, e1)) | (e1 = op(e1, e1)) | (e1 = op(e0, e1))), inference(cnf_transformation, [], [f2])).
fof(f880, plain, (spl66_47 | spl66_43 | spl66_39 | spl66_35), inference(avatar_split_clause, [], [f215, f734, f751, f768, f785])).
fof(f215, plain, ((e2 = op(e1, e3)) | (e2 = op(e1, e2)) | (e2 = op(e1, e1)) | (e2 = op(e1, e0))), inference(cnf_transformation, [], [f2])).
fof(f879, plain, (spl66_59 | spl66_43 | spl66_27 | spl66_11), inference(avatar_split_clause, [], [f216, f632, f700, f768, f836])).
fof(f216, plain, ((e2 = op(e3, e1)) | (e2 = op(e2, e1)) | (e2 = op(e1, e1)) | (e2 = op(e0, e1))), inference(cnf_transformation, [], [f2])).
fof(f878, plain, (spl66_48 | spl66_44 | spl66_40 | spl66_36), inference(avatar_split_clause, [], [f217, f738, f755, f772, f789])).
fof(f217, plain, ((e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, e0))), inference(cnf_transformation, [], [f2])).
fof(f877, plain, (spl66_60 | spl66_44 | spl66_28 | spl66_12), inference(avatar_split_clause, [], [f218, f636, f704, f772, f840])).
fof(f218, plain, ((e3 = op(e3, e1)) | (e3 = op(e2, e1)) | (e3 = op(e1, e1)) | (e3 = op(e0, e1))), inference(cnf_transformation, [], [f2])).
fof(f876, plain, (spl66_29 | spl66_25 | spl66_21 | spl66_17), inference(avatar_split_clause, [], [f219, f658, f675, f692, f709])).
fof(f219, plain, ((e0 = op(e2, e3)) | (e0 = op(e2, e2)) | (e0 = op(e2, e1)) | (e0 = op(e2, e0))), inference(cnf_transformation, [], [f2])).
fof(f875, plain, (spl66_53 | spl66_37 | spl66_21 | spl66_5), inference(avatar_split_clause, [], [f220, f607, f675, f743, f811])).
fof(f220, plain, ((e0 = op(e3, e2)) | (e0 = op(e2, e2)) | (e0 = op(e1, e2)) | (e0 = op(e0, e2))), inference(cnf_transformation, [], [f2])).
fof(f874, plain, (spl66_30 | spl66_26 | spl66_22 | spl66_18), inference(avatar_split_clause, [], [f221, f662, f679, f696, f713])).
fof(f221, plain, ((e1 = op(e2, e3)) | (e1 = op(e2, e2)) | (e1 = op(e2, e1)) | (e1 = op(e2, e0))), inference(cnf_transformation, [], [f2])).
fof(f873, plain, (spl66_54 | spl66_38 | spl66_22 | spl66_6), inference(avatar_split_clause, [], [f222, f611, f679, f747, f815])).
fof(f222, plain, ((e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(e0, e2))), inference(cnf_transformation, [], [f2])).
fof(f872, plain, (spl66_31 | spl66_27 | spl66_23 | spl66_19), inference(avatar_split_clause, [], [f223, f666, f683, f700, f717])).
fof(f223, plain, ((e2 = op(e2, e3)) | (e2 = op(e2, e2)) | (e2 = op(e2, e1)) | (e2 = op(e2, e0))), inference(cnf_transformation, [], [f2])).
fof(f871, plain, (spl66_55 | spl66_39 | spl66_23 | spl66_7), inference(avatar_split_clause, [], [f224, f615, f683, f751, f819])).
fof(f224, plain, ((e2 = op(e3, e2)) | (e2 = op(e2, e2)) | (e2 = op(e1, e2)) | (e2 = op(e0, e2))), inference(cnf_transformation, [], [f2])).
fof(f870, plain, (spl66_32 | spl66_28 | spl66_24 | spl66_20), inference(avatar_split_clause, [], [f225, f670, f687, f704, f721])).
fof(f225, plain, ((e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, e0))), inference(cnf_transformation, [], [f2])).
fof(f869, plain, (spl66_56 | spl66_40 | spl66_24 | spl66_8), inference(avatar_split_clause, [], [f226, f619, f687, f755, f823])).
fof(f226, plain, ((e3 = op(e3, e2)) | (e3 = op(e2, e2)) | (e3 = op(e1, e2)) | (e3 = op(e0, e2))), inference(cnf_transformation, [], [f2])).
fof(f868, plain, (spl66_13 | spl66_9 | spl66_5 | spl66_1), inference(avatar_split_clause, [], [f227, f590, f607, f624, f641])).
fof(f227, plain, ((e0 = op(e3, e3)) | (e0 = op(e3, e2)) | (e0 = op(e3, e1)) | (e0 = op(e3, e0))), inference(cnf_transformation, [], [f2])).
fof(f867, plain, (spl66_49 | spl66_33 | spl66_17 | spl66_1), inference(avatar_split_clause, [], [f228, f590, f658, f726, f794])).
fof(f228, plain, ((e0 = op(e3, e3)) | (e0 = op(e2, e3)) | (e0 = op(e1, e3)) | (e0 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f866, plain, (spl66_14 | spl66_10 | spl66_6 | spl66_2), inference(avatar_split_clause, [], [f229, f594, f611, f628, f645])).
fof(f229, plain, ((e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, e0))), inference(cnf_transformation, [], [f2])).
fof(f865, plain, (spl66_50 | spl66_34 | spl66_18 | spl66_2), inference(avatar_split_clause, [], [f230, f594, f662, f730, f798])).
fof(f230, plain, ((e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f864, plain, (spl66_15 | spl66_11 | spl66_7 | spl66_3), inference(avatar_split_clause, [], [f231, f598, f615, f632, f649])).
fof(f231, plain, ((e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, e0))), inference(cnf_transformation, [], [f2])).
fof(f863, plain, (spl66_51 | spl66_35 | spl66_19 | spl66_3), inference(avatar_split_clause, [], [f232, f598, f666, f734, f802])).
fof(f232, plain, ((e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f862, plain, (spl66_16 | spl66_12 | spl66_8 | spl66_4), inference(avatar_split_clause, [], [f233, f602, f619, f636, f653])).
fof(f233, plain, ((e3 = op(e3, e3)) | (e3 = op(e3, e2)) | (e3 = op(e3, e1)) | (e3 = op(e3, e0))), inference(cnf_transformation, [], [f2])).
fof(f861, plain, (spl66_52 | spl66_36 | spl66_20 | spl66_4), inference(avatar_split_clause, [], [f234, f602, f670, f738, f806])).
fof(f234, plain, ((e3 = op(e3, e3)) | (e3 = op(e2, e3)) | (e3 = op(e1, e3)) | (e3 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f860, plain, (spl66_61 | spl66_62 | spl66_63 | spl66_64), inference(avatar_split_clause, [], [f187, f857, f853, f849, f845])).
fof(f187, plain, ((op(e0, e0) = e3) | (op(e0, e0) = e2) | (op(e0, e0) = e1) | (e0 = op(e0, e0))), inference(cnf_transformation, [], [f1])).
fof(f1, plain, (((e3 = op(e3, e3)) | (e2 = op(e3, e3)) | (e1 = op(e3, e3)) | (e0 = op(e3, e3))) & ((e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (e0 = op(e3, e2))) & ((e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (e0 = op(e3, e1))) & ((e3 = op(e3, e0)) | (e2 = op(e3, e0)) | (e1 = op(e3, e0)) | (e0 = op(e3, e0))) & ((e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))) & ((e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (e0 = op(e2, e2))) & ((e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))) & ((e3 = op(e2, e0)) | (e2 = op(e2, e0)) | (e1 = op(e2, e0)) | (e0 = op(e2, e0))) & ((e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (e0 = op(e1, e3))) & ((e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (e0 = op(e1, e2))) & ((e3 = op(e1, e1)) | (e2 = op(e1, e1)) | (e1 = op(e1, e1)) | (e0 = op(e1, e1))) & ((e3 = op(e1, e0)) | (e2 = op(e1, e0)) | (e1 = op(e1, e0)) | (e0 = op(e1, e0))) & ((e3 = op(e0, e3)) | (e2 = op(e0, e3)) | (e1 = op(e0, e3)) | (e0 = op(e0, e3))) & ((e3 = op(e0, e2)) | (e2 = op(e0, e2)) | (e1 = op(e0, e2)) | (e0 = op(e0, e2))) & ((e3 = op(e0, e1)) | (e2 = op(e0, e1)) | (e1 = op(e0, e1)) | (e0 = op(e0, e1))) & ((op(e0, e0) = e3) | (op(e0, e0) = e2) | (op(e0, e0) = e1) | (e0 = op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG140+1.p', ax1)).
fof(f843, plain, (spl66_57 | spl66_58 | spl66_59 | spl66_60), inference(avatar_split_clause, [], [f188, f840, f836, f832, f828])).
fof(f188, plain, ((e3 = op(e0, e1)) | (e2 = op(e0, e1)) | (e1 = op(e0, e1)) | (e0 = op(e0, e1))), inference(cnf_transformation, [], [f1])).
fof(f826, plain, (spl66_53 | spl66_54 | spl66_55 | spl66_56), inference(avatar_split_clause, [], [f189, f823, f819, f815, f811])).
fof(f189, plain, ((e3 = op(e0, e2)) | (e2 = op(e0, e2)) | (e1 = op(e0, e2)) | (e0 = op(e0, e2))), inference(cnf_transformation, [], [f1])).
fof(f809, plain, (spl66_49 | spl66_50 | spl66_51 | spl66_52), inference(avatar_split_clause, [], [f190, f806, f802, f798, f794])).
fof(f190, plain, ((e3 = op(e0, e3)) | (e2 = op(e0, e3)) | (e1 = op(e0, e3)) | (e0 = op(e0, e3))), inference(cnf_transformation, [], [f1])).
fof(f792, plain, (spl66_45 | spl66_46 | spl66_47 | spl66_48), inference(avatar_split_clause, [], [f191, f789, f785, f781, f777])).
fof(f191, plain, ((e3 = op(e1, e0)) | (e2 = op(e1, e0)) | (e1 = op(e1, e0)) | (e0 = op(e1, e0))), inference(cnf_transformation, [], [f1])).
fof(f758, plain, (spl66_37 | spl66_38 | spl66_39 | spl66_40), inference(avatar_split_clause, [], [f193, f755, f751, f747, f743])).
fof(f193, plain, ((e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (e0 = op(e1, e2))), inference(cnf_transformation, [], [f1])).
fof(f741, plain, (spl66_33 | spl66_34 | spl66_35 | spl66_36), inference(avatar_split_clause, [], [f194, f738, f734, f730, f726])).
fof(f194, plain, ((e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (e0 = op(e1, e3))), inference(cnf_transformation, [], [f1])).
fof(f724, plain, (spl66_29 | spl66_30 | spl66_31 | spl66_32), inference(avatar_split_clause, [], [f195, f721, f717, f713, f709])).
fof(f195, plain, ((e3 = op(e2, e0)) | (e2 = op(e2, e0)) | (e1 = op(e2, e0)) | (e0 = op(e2, e0))), inference(cnf_transformation, [], [f1])).
fof(f707, plain, (spl66_25 | spl66_26 | spl66_27 | spl66_28), inference(avatar_split_clause, [], [f196, f704, f700, f696, f692])).
fof(f196, plain, ((e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))), inference(cnf_transformation, [], [f1])).
fof(f690, plain, (spl66_21 | spl66_22 | spl66_23 | spl66_24), inference(avatar_split_clause, [], [f197, f687, f683, f679, f675])).
fof(f197, plain, ((e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (e0 = op(e2, e2))), inference(cnf_transformation, [], [f1])).
fof(f673, plain, (spl66_17 | spl66_18 | spl66_19 | spl66_20), inference(avatar_split_clause, [], [f198, f670, f666, f662, f658])).
fof(f198, plain, ((e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))), inference(cnf_transformation, [], [f1])).
fof(f656, plain, (spl66_13 | spl66_14 | spl66_15 | spl66_16), inference(avatar_split_clause, [], [f199, f653, f649, f645, f641])).
fof(f199, plain, ((e3 = op(e3, e0)) | (e2 = op(e3, e0)) | (e1 = op(e3, e0)) | (e0 = op(e3, e0))), inference(cnf_transformation, [], [f1])).
fof(f639, plain, (spl66_9 | spl66_10 | spl66_11 | spl66_12), inference(avatar_split_clause, [], [f200, f636, f632, f628, f624])).
fof(f200, plain, ((e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (e0 = op(e3, e1))), inference(cnf_transformation, [], [f1])).
fof(f622, plain, (spl66_5 | spl66_6 | spl66_7 | spl66_8), inference(avatar_split_clause, [], [f201, f619, f615, f611, f607])).
fof(f201, plain, ((e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (e0 = op(e3, e2))), inference(cnf_transformation, [], [f1])).
fof(f605, plain, (spl66_1 | spl66_2 | spl66_3 | spl66_4), inference(avatar_split_clause, [], [f202, f602, f598, f594, f590])).
fof(f202, plain, ((e3 = op(e3, e3)) | (e2 = op(e3, e3)) | (e1 = op(e3, e3)) | (e0 = op(e3, e3))), inference(cnf_transformation, [], [f1])).