fof(f4038, plain, $false, inference(avatar_sat_refutation, [], [f327, f344, f361, f378, f395, f412, f429, f446, f463, f480, f497, f514, f531, f548, f565, f582, f583, f584, f585, f586, f587, f588, f589, f590, f591, f592, f593, f594, f595, f596, f597, f598, f599, f600, f601, f602, f603, f604, f605, f606, f607, f608, f609, f610, f611, f612, f613, f614, f619, f620, f621, f622, f623, f624, f629, f630, f631, f632, f633, f634, f639, f640, f641, f642, f643, f644, f649, f650, f651, f652, f653, f654, f660, f664, f669, f670, f671, f672, f673, f674, f679, f680, f681, f682, f683, f684, f689, f690, f691, f692, f693, f694, f699, f700, f701, f702, f703, f704, f711, f714, f719, f720, f721, f722, f723, f724, f729, f730, f731, f732, f733, f734, f740, f741, f742, f743, f744, f749, f750, f751, f752, f753, f754, f762, f764, f765, f766, f770, f787, f804, f821, f838, f847, f856, f861, f870, f875, f880, f885, f890, f895, f904, f909, f919, f930, f935, f940, f945, f950, f951, f956, f957, f984, f985, f987, f996, f997, f1004, f1007, f1015, f1028, f1044, f1047, f1061, f1069, f1070, f1072, f1074, f1077, f1080, f1083, f1092, f1096, f1099, f1102, f1107, f1110, f1113, f1121, f1122, f1137, f1139, f1145, f1157, f1165, f1174, f1177, f1186, f1187, f1190, f1193, f1221, f1236, f1249, f1251, f1256, f1263, f1267, f1269, f1271, f1276, f1278, f1292, f1297, f1305, f1314, f1328, f1349, f1352, f1356, f1371, f1373, f1377, f1379, f1389, f1392, f1396, f1399, f1420, f1432, f1439, f1450, f1463, f1492, f1497, f1523, f1541, f1570, f1576, f1590, f1596, f1601, f1620, f1625, f1637, f1655, f1657, f1665, f1678, f1679, f1719, f1725, f1727, f1751, f1756, f1761, f1778, f1789, f1804, f1814, f1824, f1844, f1846, f1862, f1885, f1893, f1897, f1903, f1904, f1916, f1919, f1920, f1921, f1922, f1927, f1944, f1948, f1953, f1961, f1962, f1975, f1977, f1984, f2003, f2036, f2044, f2075, f2085, f2106, f2118, f2120, f2148, f2157, f2172, f2175, f2184, f2190, f2196, f2204, f2208, f2216, f2228, f2233, f2243, f2246, f2267, f2294, f2306, f2312, f2313, f2323, f2333, f2337, f2344, f2367, f2396, f2401, f2423, f2428, f2432, f2437, f2439, f2459, f2489, f2494, f2512, f2519, f2533, f2537, f2549, f2557, f2573, f2610, f2615, f2631, f2670, f2684, f2702, f2717, f2719, f2729, f2735, f2738, f2743, f2744, f2750, f2771, f2779, f2790, f2805, f2821, f2826, f2835, f2841, f2859, f2871, f2878, f2896, f2900, f2911, f2914, f2950, f2960, f2972, f2978, f2979, f2986, f2994, f3001, f3013, f3022, f3029, f3035, f3067, f3069, f3074, f3082, f3090, f3099, f3101, f3109, f3114, f3122, f3131, f3150, f3156, f3157, f3171, f3209, f3210, f3211, f3215, f3227, f3228, f3232, f3244, f3250, f3265, f3269, f3272, f3274, f3289, f3294, f3295, f3300, f3311, f3312, f3325, f3327, f3336, f3337, f3345, f3350, f3351, f3353, f3360, f3361, f3392, f3395, f3406, f3409, f3430, f3431, f3433, f3445, f3446, f3455, f3464, f3465, f3470, f3475, f3483, f3485, f3487, f3493, f3502, f3513, f3519, f3528, f3535, f3536, f3543, f3548, f3556, f3557, f3564, f3570, f3578, f3584, f3589, f3592, f3595, f3597, f3599, f3601, f3603, f3609, f3616, f3619, f3628, f3638, f3646, f3667, f3672, f3676, f3687, f3688, f3689, f3698, f3699, f3706, f3711, f3712, f3718, f3719, f3720, f3729, f3731, f3737, f3739, f3748, f3754, f3755, f3764, f3768, f3774, f3805, f3813, f3823, f3824, f3833, f3839, f3848, f3849, f3854, f3865, f3866, f3871, f3880, f3890, f3893, f3900, f3925, f3947, f3948, f3976, f3977, f3978, f3983, f3991, f4002, f4010, f4014, f4027, f4031])).
fof(f4031, plain, (~ spl15_4 | ~ spl15_8), inference(avatar_split_clause, [], [f4029, f341, f324])).
fof(f324, plain, (spl15_4 <=> (e3 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl15_4])])).
fof(f341, plain, (spl15_8 <=> (e3 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl15_8])])).
fof(f4029, plain, (~ (e3 = op(e3, e3)) | ~ spl15_8), inference(backward_demodulation, [], [f180, f343])).
fof(f343, plain, ((e3 = op(e3, e2)) | ~ spl15_8), inference(avatar_component_clause, [], [f341])).
fof(f180, plain, ~ (op(e3, e2) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f3, plain, (~ (op(e3, e2) = op(e3, e3)) & ~ (op(e3, e1) = op(e3, e3)) & ~ (op(e3, e1) = op(e3, e2)) & ~ (op(e3, e0) = op(e3, e3)) & ~ (op(e3, e0) = op(e3, e2)) & ~ (op(e3, e0) = op(e3, e1)) & ~ (op(e2, e2) = op(e2, e3)) & ~ (op(e2, e1) = op(e2, e3)) & ~ (op(e2, e1) = op(e2, e2)) & ~ (op(e2, e0) = op(e2, e3)) & ~ (op(e2, e0) = op(e2, e2)) & ~ (op(e2, e0) = op(e2, e1)) & ~ (op(e1, e2) = op(e1, e3)) & ~ (op(e1, e1) = op(e1, e3)) & ~ (op(e1, e1) = op(e1, e2)) & ~ (op(e1, e0) = op(e1, e3)) & ~ (op(e1, e0) = op(e1, e2)) & ~ (op(e1, e0) = op(e1, e1)) & ~ (op(e0, e2) = op(e0, e3)) & ~ (op(e0, e1) = op(e0, e3)) & ~ (op(e0, e1) = op(e0, e2)) & ~ (op(e0, e0) = op(e0, e3)) & ~ (op(e0, e0) = op(e0, e2)) & ~ (op(e0, e0) = op(e0, e1)) & ~ (op(e2, e3) = op(e3, e3)) & ~ (op(e1, e3) = op(e3, e3)) & ~ (op(e1, e3) = op(e2, e3)) & ~ (op(e0, e3) = op(e3, e3)) & ~ (op(e0, e3) = op(e2, e3)) & ~ (op(e0, e3) = op(e1, e3)) & ~ (op(e2, e2) = op(e3, e2)) & ~ (op(e1, e2) = op(e3, e2)) & ~ (op(e1, e2) = op(e2, e2)) & ~ (op(e0, e2) = op(e3, e2)) & ~ (op(e0, e2) = op(e2, e2)) & ~ (op(e0, e2) = op(e1, e2)) & ~ (op(e2, e1) = op(e3, e1)) & ~ (op(e1, e1) = op(e3, e1)) & ~ (op(e1, e1) = op(e2, e1)) & ~ (op(e0, e1) = op(e3, e1)) & ~ (op(e0, e1) = op(e2, e1)) & ~ (op(e0, e1) = op(e1, e1)) & ~ (op(e2, e0) = op(e3, e0)) & ~ (op(e1, e0) = op(e3, e0)) & ~ (op(e1, e0) = op(e2, e0)) & ~ (op(e0, e0) = op(e3, e0)) & ~ (op(e0, e0) = op(e2, e0)) & ~ (op(e0, e0) = op(e1, e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG131+1.p', ax3)).
fof(f4027, plain, (~ spl15_5 | ~ spl15_13), inference(avatar_split_clause, [], [f4021, f363, f329])).
fof(f329, plain, (spl15_5 <=> (e0 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl15_5])])).
fof(f363, plain, (spl15_13 <=> (e0 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl15_13])])).
fof(f4021, plain, (~ (e0 = op(e3, e2)) | ~ spl15_13), inference(backward_demodulation, [], [f176, f365])).
fof(f365, plain, ((e0 = op(e3, e0)) | ~ spl15_13), inference(avatar_component_clause, [], [f363])).
fof(f176, plain, ~ (op(e3, e0) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f4014, plain, (~ spl15_24 | ~ spl15_32), inference(avatar_split_clause, [], [f4012, f443, f409])).
fof(f409, plain, (spl15_24 <=> (e3 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl15_24])])).
fof(f443, plain, (spl15_32 <=> (e3 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl15_32])])).
fof(f4012, plain, (~ (e3 = op(e2, e2)) | ~ spl15_32), inference(backward_demodulation, [], [f170, f445])).
fof(f445, plain, ((e3 = op(e2, e0)) | ~ spl15_32), inference(avatar_component_clause, [], [f443])).
fof(f170, plain, ~ (op(e2, e0) = op(e2, e2)), inference(cnf_transformation, [], [f3])).
fof(f4010, plain, (~ spl15_34 | ~ spl15_36), inference(avatar_contradiction_clause, [], [f4009])).
fof(f4009, plain, ($false | (~ spl15_34 | ~ spl15_36)), inference(subsumption_resolution, [], [f4008, f185])).
fof(f185, plain, ~ (e1 = e3), inference(cnf_transformation, [], [f4])).
fof(f4, plain, (~ (e2 = e3) & ~ (e1 = e3) & ~ (e1 = e2) & ~ (e0 = e3) & ~ (e0 = e2) & ~ (e0 = e1)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG131+1.p', ax4)).
fof(f4008, plain, ((e1 = e3) | (~ spl15_34 | ~ spl15_36)), inference(forward_demodulation, [], [f462, f454])).
fof(f454, plain, ((e1 = op(e1, e3)) | ~ spl15_34), inference(avatar_component_clause, [], [f452])).
fof(f452, plain, (spl15_34 <=> (e1 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl15_34])])).
fof(f462, plain, ((e3 = op(e1, e3)) | ~ spl15_36), inference(avatar_component_clause, [], [f460])).
fof(f460, plain, (spl15_36 <=> (e3 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl15_36])])).
fof(f4002, plain, (~ spl15_30 | ~ spl15_46), inference(avatar_split_clause, [], [f3996, f503, f435])).
fof(f435, plain, (spl15_30 <=> (e1 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl15_30])])).
fof(f503, plain, (spl15_46 <=> (e1 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl15_46])])).
fof(f3996, plain, (~ (e1 = op(e2, e0)) | ~ spl15_46), inference(backward_demodulation, [], [f136, f505])).
fof(f505, plain, ((e1 = op(e1, e0)) | ~ spl15_46), inference(avatar_component_clause, [], [f503])).
fof(f136, plain, ~ (op(e1, e0) = op(e2, e0)), inference(cnf_transformation, [], [f3])).
fof(f3991, plain, (~ spl15_5 | ~ spl15_53), inference(avatar_split_clause, [], [f3985, f533, f329])).
fof(f533, plain, (spl15_53 <=> (e0 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl15_53])])).
fof(f3985, plain, (~ (e0 = op(e3, e2)) | ~ spl15_53), inference(backward_demodulation, [], [f147, f535])).
fof(f535, plain, ((e0 = op(e0, e2)) | ~ spl15_53), inference(avatar_component_clause, [], [f533])).
fof(f147, plain, ~ (op(e0, e2) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f3983, plain, (~ spl15_44 | ~ spl15_60), inference(avatar_split_clause, [], [f3979, f562, f494])).
fof(f494, plain, (spl15_44 <=> (e3 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl15_44])])).
fof(f562, plain, (spl15_60 <=> (e3 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl15_60])])).
fof(f3979, plain, (~ (e3 = op(e1, e1)) | ~ spl15_60), inference(backward_demodulation, [], [f139, f564])).
fof(f564, plain, ((e3 = op(e0, e1)) | ~ spl15_60), inference(avatar_component_clause, [], [f562])).
fof(f139, plain, ~ (op(e0, e1) = op(e1, e1)), inference(cnf_transformation, [], [f3])).
fof(f3978, plain, (spl15_53 | ~ spl15_63 | ~ spl15_92), inference(avatar_split_clause, [], [f3975, f823, f575, f533])).
fof(f575, plain, (spl15_63 <=> (op(e0, e0) = e2)), introduced(avatar_definition, [new_symbols(naming, [spl15_63])])).
fof(f823, plain, (spl15_92 <=> (e0 = op(e0, op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl15_92])])).
fof(f3975, plain, ((e0 = op(e0, e2)) | (~ spl15_63 | ~ spl15_92)), inference(backward_demodulation, [], [f825, f577])).
fof(f577, plain, ((op(e0, e0) = e2) | ~ spl15_63), inference(avatar_component_clause, [], [f575])).
fof(f825, plain, ((e0 = op(e0, op(e0, e0))) | ~ spl15_92), inference(avatar_component_clause, [], [f823])).
fof(f3977, plain, (~ spl15_51 | ~ spl15_63), inference(avatar_split_clause, [], [f3973, f575, f524])).
fof(f524, plain, (spl15_51 <=> (e2 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl15_51])])).
fof(f3973, plain, (~ (e2 = op(e0, e3)) | ~ spl15_63), inference(backward_demodulation, [], [f159, f577])).
fof(f159, plain, ~ (op(e0, e0) = op(e0, e3)), inference(cnf_transformation, [], [f3])).
fof(f3976, plain, (~ spl15_15 | ~ spl15_63), inference(avatar_split_clause, [], [f3970, f575, f371])).
fof(f371, plain, (spl15_15 <=> (e2 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl15_15])])).
fof(f3970, plain, (~ (e2 = op(e3, e0)) | ~ spl15_63), inference(backward_demodulation, [], [f135, f577])).
fof(f135, plain, ~ (op(e0, e0) = op(e3, e0)), inference(cnf_transformation, [], [f3])).
fof(f3948, plain, (~ spl15_50 | ~ spl15_34), inference(avatar_split_clause, [], [f3798, f452, f520])).
fof(f520, plain, (spl15_50 <=> (e1 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl15_50])])).
fof(f3798, plain, (~ (e1 = op(e0, e3)) | ~ spl15_34), inference(forward_demodulation, [], [f151, f454])).
fof(f151, plain, ~ (op(e0, e3) = op(e1, e3)), inference(cnf_transformation, [], [f3])).
fof(f3947, plain, (~ spl15_46 | ~ spl15_34), inference(avatar_split_clause, [], [f3796, f452, f503])).
fof(f3796, plain, (~ (e1 = op(e1, e0)) | ~ spl15_34), inference(forward_demodulation, [], [f165, f454])).
fof(f165, plain, ~ (op(e1, e0) = op(e1, e3)), inference(cnf_transformation, [], [f3])).
fof(f3925, plain, (~ spl15_17 | ~ spl15_24 | spl15_110), inference(avatar_contradiction_clause, [], [f3924])).
fof(f3924, plain, ($false | (~ spl15_17 | ~ spl15_24 | spl15_110)), inference(subsumption_resolution, [], [f3923, f382])).
fof(f382, plain, ((e0 = op(e2, e3)) | ~ spl15_17), inference(avatar_component_clause, [], [f380])).
fof(f380, plain, (spl15_17 <=> (e0 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl15_17])])).
fof(f3923, plain, (~ (e0 = op(e2, e3)) | (~ spl15_24 | spl15_110)), inference(forward_demodulation, [], [f908, f411])).
fof(f411, plain, ((e3 = op(e2, e2)) | ~ spl15_24), inference(avatar_component_clause, [], [f409])).
fof(f908, plain, (~ (e0 = op(e2, op(e2, e2))) | spl15_110), inference(avatar_component_clause, [], [f906])).
fof(f906, plain, (spl15_110 <=> (e0 = op(e2, op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl15_110])])).
fof(f3900, plain, (~ spl15_1 | ~ spl15_5), inference(avatar_split_clause, [], [f3895, f329, f312])).
fof(f312, plain, (spl15_1 <=> (e0 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl15_1])])).
fof(f3895, plain, (~ (e0 = op(e3, e3)) | ~ spl15_5), inference(backward_demodulation, [], [f180, f331])).
fof(f331, plain, ((e0 = op(e3, e2)) | ~ spl15_5), inference(avatar_component_clause, [], [f329])).
fof(f3893, plain, (~ spl15_17 | ~ spl15_24 | ~ spl15_54 | spl15_106), inference(avatar_contradiction_clause, [], [f3892])).
fof(f3892, plain, ($false | (~ spl15_17 | ~ spl15_24 | ~ spl15_54 | spl15_106)), inference(subsumption_resolution, [], [f3888, f539])).
fof(f539, plain, ((e1 = op(e0, e2)) | ~ spl15_54), inference(avatar_component_clause, [], [f537])).
fof(f537, plain, (spl15_54 <=> (e1 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl15_54])])).
fof(f3888, plain, (~ (e1 = op(e0, e2)) | (~ spl15_17 | ~ spl15_24 | spl15_106)), inference(backward_demodulation, [], [f3878, f382])).
fof(f3878, plain, (~ (e1 = op(op(e2, e3), e2)) | (~ spl15_24 | spl15_106)), inference(backward_demodulation, [], [f889, f411])).
fof(f889, plain, (~ (e1 = op(op(e2, op(e2, e2)), e2)) | spl15_106), inference(avatar_component_clause, [], [f887])).
fof(f887, plain, (spl15_106 <=> (e1 = op(op(e2, op(e2, e2)), e2))), introduced(avatar_definition, [new_symbols(naming, [spl15_106])])).
fof(f3890, plain, (~ spl15_1 | ~ spl15_17), inference(avatar_split_clause, [], [f3882, f380, f312])).
fof(f3882, plain, (~ (e0 = op(e3, e3)) | ~ spl15_17), inference(backward_demodulation, [], [f156, f382])).
fof(f156, plain, ~ (op(e2, e3) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f3880, plain, (~ spl15_8 | ~ spl15_24), inference(avatar_split_clause, [], [f3872, f409, f341])).
fof(f3872, plain, (~ (e3 = op(e3, e2)) | ~ spl15_24), inference(backward_demodulation, [], [f150, f411])).
fof(f150, plain, ~ (op(e2, e2) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f3871, plain, (~ spl15_19 | ~ spl15_27), inference(avatar_split_clause, [], [f3868, f422, f388])).
fof(f388, plain, (spl15_19 <=> (e2 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl15_19])])).
fof(f422, plain, (spl15_27 <=> (e2 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl15_27])])).
fof(f3868, plain, (~ (e2 = op(e2, e3)) | ~ spl15_27), inference(backward_demodulation, [], [f173, f424])).
fof(f424, plain, ((e2 = op(e2, e1)) | ~ spl15_27), inference(avatar_component_clause, [], [f422])).
fof(f173, plain, ~ (op(e2, e1) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f3866, plain, (spl15_27 | ~ spl15_30 | ~ spl15_94), inference(avatar_split_clause, [], [f3864, f831, f435, f422])).
fof(f831, plain, (spl15_94 <=> (e2 = op(e2, op(e2, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl15_94])])).
fof(f3864, plain, ((e2 = op(e2, e1)) | (~ spl15_30 | ~ spl15_94)), inference(backward_demodulation, [], [f833, f437])).
fof(f437, plain, ((e1 = op(e2, e0)) | ~ spl15_30), inference(avatar_component_clause, [], [f435])).
fof(f833, plain, ((e2 = op(e2, op(e2, e0))) | ~ spl15_94), inference(avatar_component_clause, [], [f831])).
fof(f3865, plain, (~ spl15_22 | ~ spl15_30), inference(avatar_split_clause, [], [f3860, f435, f401])).
fof(f401, plain, (spl15_22 <=> (e1 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl15_22])])).
fof(f3860, plain, (~ (e1 = op(e2, e2)) | ~ spl15_30), inference(backward_demodulation, [], [f170, f437])).
fof(f3854, plain, (~ spl15_28 | ~ spl15_44), inference(avatar_split_clause, [], [f3851, f494, f426])).
fof(f426, plain, (spl15_28 <=> (e3 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl15_28])])).
fof(f3851, plain, (~ (e3 = op(e2, e1)) | ~ spl15_44), inference(backward_demodulation, [], [f142, f496])).
fof(f496, plain, ((e3 = op(e1, e1)) | ~ spl15_44), inference(avatar_component_clause, [], [f494])).
fof(f142, plain, ~ (op(e1, e1) = op(e2, e1)), inference(cnf_transformation, [], [f3])).
fof(f3849, plain, (~ spl15_37 | ~ spl15_45), inference(avatar_split_clause, [], [f3843, f499, f465])).
fof(f465, plain, (spl15_37 <=> (e0 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl15_37])])).
fof(f499, plain, (spl15_45 <=> (e0 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl15_45])])).
fof(f3843, plain, (~ (e0 = op(e1, e2)) | ~ spl15_45), inference(backward_demodulation, [], [f164, f501])).
fof(f501, plain, ((e0 = op(e1, e0)) | ~ spl15_45), inference(avatar_component_clause, [], [f499])).
fof(f164, plain, ~ (op(e1, e0) = op(e1, e2)), inference(cnf_transformation, [], [f3])).
fof(f3848, plain, (~ spl15_29 | ~ spl15_45), inference(avatar_split_clause, [], [f3841, f499, f431])).
fof(f431, plain, (spl15_29 <=> (e0 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl15_29])])).
fof(f3841, plain, (~ (e0 = op(e2, e0)) | ~ spl15_45), inference(backward_demodulation, [], [f136, f501])).
fof(f3839, plain, (~ spl15_19 | ~ spl15_51), inference(avatar_split_clause, [], [f3834, f524, f388])).
fof(f3834, plain, (~ (e2 = op(e2, e3)) | ~ spl15_51), inference(backward_demodulation, [], [f152, f526])).
fof(f526, plain, ((e2 = op(e0, e3)) | ~ spl15_51), inference(avatar_component_clause, [], [f524])).
fof(f152, plain, ~ (op(e0, e3) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f3833, plain, (~ spl15_22 | ~ spl15_54), inference(avatar_split_clause, [], [f3827, f537, f401])).
fof(f3827, plain, (~ (e1 = op(e2, e2)) | ~ spl15_54), inference(backward_demodulation, [], [f146, f539])).
fof(f146, plain, ~ (op(e0, e2) = op(e2, e2)), inference(cnf_transformation, [], [f3])).
fof(f3824, plain, (~ spl15_52 | ~ spl15_64), inference(avatar_split_clause, [], [f3818, f579, f528])).
fof(f528, plain, (spl15_52 <=> (e3 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl15_52])])).
fof(f579, plain, (spl15_64 <=> (op(e0, e0) = e3)), introduced(avatar_definition, [new_symbols(naming, [spl15_64])])).
fof(f3818, plain, (~ (e3 = op(e0, e3)) | ~ spl15_64), inference(backward_demodulation, [], [f159, f581])).
fof(f581, plain, ((op(e0, e0) = e3) | ~ spl15_64), inference(avatar_component_clause, [], [f579])).
fof(f3823, plain, (~ spl15_48 | ~ spl15_64), inference(avatar_split_clause, [], [f3815, f579, f511])).
fof(f511, plain, (spl15_48 <=> (e3 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl15_48])])).
fof(f3815, plain, (~ (e3 = op(e1, e0)) | ~ spl15_64), inference(backward_demodulation, [], [f133, f581])).
fof(f133, plain, ~ (op(e0, e0) = op(e1, e0)), inference(cnf_transformation, [], [f3])).
fof(f3813, plain, (~ spl15_1 | ~ spl15_83 | spl15_96), inference(avatar_split_clause, [], [f3812, f840, f784, f312])).
fof(f784, plain, (spl15_83 <=> (e3 = op(e3, op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl15_83])])).
fof(f840, plain, (spl15_96 <=> (e0 = op(op(e3, op(e3, e3)), e3))), introduced(avatar_definition, [new_symbols(naming, [spl15_96])])).
fof(f3812, plain, (~ (e0 = op(e3, e3)) | (~ spl15_83 | spl15_96)), inference(backward_demodulation, [], [f842, f786])).
fof(f786, plain, ((e3 = op(e3, op(e3, e3))) | ~ spl15_83), inference(avatar_component_clause, [], [f784])).
fof(f842, plain, (~ (e0 = op(op(e3, op(e3, e3)), e3)) | spl15_96), inference(avatar_component_clause, [], [f840])).
fof(f3805, plain, (~ spl15_8 | ~ spl15_15 | spl15_95), inference(avatar_split_clause, [], [f3804, f835, f371, f341])).
fof(f835, plain, (spl15_95 <=> (e3 = op(e3, op(e3, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl15_95])])).
fof(f3804, plain, (~ (e3 = op(e3, e2)) | (~ spl15_15 | spl15_95)), inference(forward_demodulation, [], [f836, f373])).
fof(f373, plain, ((e2 = op(e3, e0)) | ~ spl15_15), inference(avatar_component_clause, [], [f371])).
fof(f836, plain, (~ (e3 = op(e3, op(e3, e0))) | spl15_95), inference(avatar_component_clause, [], [f835])).
fof(f3774, plain, (~ spl15_3 | ~ spl15_15), inference(avatar_split_clause, [], [f3772, f371, f320])).
fof(f320, plain, (spl15_3 <=> (e2 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl15_3])])).
fof(f3772, plain, (~ (e2 = op(e3, e3)) | ~ spl15_15), inference(backward_demodulation, [], [f177, f373])).
fof(f177, plain, ~ (op(e3, e0) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f3768, plain, (~ spl15_3 | ~ spl15_19), inference(avatar_split_clause, [], [f3765, f388, f320])).
fof(f3765, plain, (~ (e2 = op(e3, e3)) | ~ spl15_19), inference(backward_demodulation, [], [f156, f390])).
fof(f390, plain, ((e2 = op(e2, e3)) | ~ spl15_19), inference(avatar_component_clause, [], [f388])).
fof(f3764, plain, (~ spl15_18 | ~ spl15_22), inference(avatar_split_clause, [], [f3757, f401, f384])).
fof(f384, plain, (spl15_18 <=> (e1 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl15_18])])).
fof(f3757, plain, (~ (e1 = op(e2, e3)) | ~ spl15_22), inference(backward_demodulation, [], [f174, f403])).
fof(f403, plain, ((e1 = op(e2, e2)) | ~ spl15_22), inference(avatar_component_clause, [], [f401])).
fof(f174, plain, ~ (op(e2, e2) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f3755, plain, (spl15_19 | ~ spl15_28 | ~ spl15_90), inference(avatar_split_clause, [], [f3753, f814, f426, f388])).
fof(f814, plain, (spl15_90 <=> (e2 = op(e2, op(e2, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl15_90])])).
fof(f3753, plain, ((e2 = op(e2, e3)) | (~ spl15_28 | ~ spl15_90)), inference(backward_demodulation, [], [f816, f428])).
fof(f428, plain, ((e3 = op(e2, e1)) | ~ spl15_28), inference(avatar_component_clause, [], [f426])).
fof(f816, plain, ((e2 = op(e2, op(e2, e1))) | ~ spl15_90), inference(avatar_component_clause, [], [f814])).
fof(f3754, plain, (~ spl15_12 | ~ spl15_28), inference(avatar_split_clause, [], [f3750, f426, f358])).
fof(f358, plain, (spl15_12 <=> (e3 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl15_12])])).
fof(f3750, plain, (~ (e3 = op(e3, e1)) | ~ spl15_28), inference(backward_demodulation, [], [f144, f428])).
fof(f144, plain, ~ (op(e2, e1) = op(e3, e1)), inference(cnf_transformation, [], [f3])).
fof(f3748, plain, (~ spl15_25 | ~ spl15_29), inference(avatar_split_clause, [], [f3741, f431, f414])).
fof(f414, plain, (spl15_25 <=> (e0 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl15_25])])).
fof(f3741, plain, (~ (e0 = op(e2, e1)) | ~ spl15_29), inference(backward_demodulation, [], [f169, f433])).
fof(f433, plain, ((e0 = op(e2, e0)) | ~ spl15_29), inference(avatar_component_clause, [], [f431])).
fof(f169, plain, ~ (op(e2, e0) = op(e2, e1)), inference(cnf_transformation, [], [f3])).
fof(f3739, plain, (~ spl15_34 | spl15_42 | ~ spl15_81), inference(avatar_contradiction_clause, [], [f3738])).
fof(f3738, plain, ($false | (~ spl15_34 | spl15_42 | ~ spl15_81)), inference(subsumption_resolution, [], [f3736, f487])).
fof(f487, plain, (~ (e1 = op(e1, e1)) | spl15_42), inference(avatar_component_clause, [], [f486])).
fof(f486, plain, (spl15_42 <=> (e1 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl15_42])])).
fof(f3736, plain, ((e1 = op(e1, e1)) | (~ spl15_34 | ~ spl15_81)), inference(backward_demodulation, [], [f778, f454])).
fof(f778, plain, ((e1 = op(e1, op(e1, e3))) | ~ spl15_81), inference(avatar_component_clause, [], [f776])).
fof(f776, plain, (spl15_81 <=> (e1 = op(e1, op(e1, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl15_81])])).
fof(f3737, plain, (~ spl15_18 | ~ spl15_34), inference(avatar_split_clause, [], [f3732, f452, f384])).
fof(f3732, plain, (~ (e1 = op(e2, e3)) | ~ spl15_34), inference(backward_demodulation, [], [f154, f454])).
fof(f154, plain, ~ (op(e1, e3) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f3731, plain, (~ spl15_37 | spl15_46 | ~ spl15_85), inference(avatar_contradiction_clause, [], [f3730])).
fof(f3730, plain, ($false | (~ spl15_37 | spl15_46 | ~ spl15_85)), inference(subsumption_resolution, [], [f3727, f504])).
fof(f504, plain, (~ (e1 = op(e1, e0)) | spl15_46), inference(avatar_component_clause, [], [f503])).
fof(f3727, plain, ((e1 = op(e1, e0)) | (~ spl15_37 | ~ spl15_85)), inference(backward_demodulation, [], [f795, f467])).
fof(f467, plain, ((e0 = op(e1, e2)) | ~ spl15_37), inference(avatar_component_clause, [], [f465])).
fof(f795, plain, ((e1 = op(e1, op(e1, e2))) | ~ spl15_85), inference(avatar_component_clause, [], [f793])).
fof(f793, plain, (spl15_85 <=> (e1 = op(e1, op(e1, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl15_85])])).
fof(f3729, plain, (~ spl15_5 | ~ spl15_37), inference(avatar_split_clause, [], [f3722, f465, f329])).
fof(f3722, plain, (~ (e0 = op(e3, e2)) | ~ spl15_37), inference(backward_demodulation, [], [f149, f467])).
fof(f149, plain, ~ (op(e1, e2) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f3720, plain, (~ spl15_34 | ~ spl15_48 | spl15_93), inference(avatar_split_clause, [], [f3717, f827, f511, f452])).
fof(f827, plain, (spl15_93 <=> (e1 = op(e1, op(e1, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl15_93])])).
fof(f3717, plain, (~ (e1 = op(e1, e3)) | (~ spl15_48 | spl15_93)), inference(backward_demodulation, [], [f828, f513])).
fof(f513, plain, ((e3 = op(e1, e0)) | ~ spl15_48), inference(avatar_component_clause, [], [f511])).
fof(f828, plain, (~ (e1 = op(e1, op(e1, e0))) | spl15_93), inference(avatar_component_clause, [], [f827])).
fof(f3719, plain, (~ spl15_36 | ~ spl15_48), inference(avatar_split_clause, [], [f3716, f511, f460])).
fof(f3716, plain, (~ (e3 = op(e1, e3)) | ~ spl15_48), inference(backward_demodulation, [], [f165, f513])).
fof(f3718, plain, (~ spl15_32 | ~ spl15_48), inference(avatar_split_clause, [], [f3713, f511, f443])).
fof(f3713, plain, (~ (e3 = op(e2, e0)) | ~ spl15_48), inference(backward_demodulation, [], [f136, f513])).
fof(f3712, plain, (~ spl15_49 | ~ spl15_52 | spl15_80), inference(avatar_split_clause, [], [f3710, f772, f528, f516])).
fof(f516, plain, (spl15_49 <=> (e0 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl15_49])])).
fof(f772, plain, (spl15_80 <=> (e0 = op(e0, op(e0, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl15_80])])).
fof(f3710, plain, (~ (e0 = op(e0, e3)) | (~ spl15_52 | spl15_80)), inference(backward_demodulation, [], [f773, f530])).
fof(f530, plain, ((e3 = op(e0, e3)) | ~ spl15_52), inference(avatar_component_clause, [], [f528])).
fof(f773, plain, (~ (e0 = op(e0, op(e0, e3))) | spl15_80), inference(avatar_component_clause, [], [f772])).
fof(f3711, plain, (~ spl15_36 | ~ spl15_52), inference(avatar_split_clause, [], [f3707, f528, f460])).
fof(f3707, plain, (~ (e3 = op(e1, e3)) | ~ spl15_52), inference(backward_demodulation, [], [f151, f530])).
fof(f3706, plain, (~ spl15_23 | ~ spl15_55), inference(avatar_split_clause, [], [f3701, f541, f405])).
fof(f405, plain, (spl15_23 <=> (e2 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl15_23])])).
fof(f541, plain, (spl15_55 <=> (e2 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl15_55])])).
fof(f3701, plain, (~ (e2 = op(e2, e2)) | ~ spl15_55), inference(backward_demodulation, [], [f146, f543])).
fof(f543, plain, ((e2 = op(e0, e2)) | ~ spl15_55), inference(avatar_component_clause, [], [f541])).
fof(f3699, plain, (~ spl15_49 | ~ spl15_57), inference(avatar_split_clause, [], [f3693, f550, f516])).
fof(f550, plain, (spl15_57 <=> (e0 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl15_57])])).
fof(f3693, plain, (~ (e0 = op(e0, e3)) | ~ spl15_57), inference(backward_demodulation, [], [f161, f552])).
fof(f552, plain, ((e0 = op(e0, e1)) | ~ spl15_57), inference(avatar_component_clause, [], [f550])).
fof(f161, plain, ~ (op(e0, e1) = op(e0, e3)), inference(cnf_transformation, [], [f3])).
fof(f3698, plain, (~ spl15_25 | ~ spl15_57), inference(avatar_split_clause, [], [f3690, f550, f414])).
fof(f3690, plain, (~ (e0 = op(e2, e1)) | ~ spl15_57), inference(backward_demodulation, [], [f140, f552])).
fof(f140, plain, ~ (op(e0, e1) = op(e2, e1)), inference(cnf_transformation, [], [f3])).
fof(f3689, plain, (spl15_57 | ~ spl15_62 | ~ spl15_92), inference(avatar_split_clause, [], [f3686, f823, f571, f550])).
fof(f571, plain, (spl15_62 <=> (op(e0, e0) = e1)), introduced(avatar_definition, [new_symbols(naming, [spl15_62])])).
fof(f3686, plain, ((e0 = op(e0, e1)) | (~ spl15_62 | ~ spl15_92)), inference(backward_demodulation, [], [f825, f573])).
fof(f573, plain, ((op(e0, e0) = e1) | ~ spl15_62), inference(avatar_component_clause, [], [f571])).
fof(f3688, plain, (~ spl15_58 | ~ spl15_62), inference(avatar_split_clause, [], [f3681, f571, f554])).
fof(f554, plain, (spl15_58 <=> (e1 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl15_58])])).
fof(f3681, plain, (~ (e1 = op(e0, e1)) | ~ spl15_62), inference(backward_demodulation, [], [f157, f573])).
fof(f157, plain, ~ (op(e0, e0) = op(e0, e1)), inference(cnf_transformation, [], [f3])).
fof(f3687, plain, (~ spl15_14 | ~ spl15_62), inference(avatar_split_clause, [], [f3680, f571, f367])).
fof(f367, plain, (spl15_14 <=> (e1 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl15_14])])).
fof(f3680, plain, (~ (e1 = op(e3, e0)) | ~ spl15_62), inference(backward_demodulation, [], [f135, f573])).
fof(f3676, plain, (spl15_62 | ~ spl15_92 | ~ spl15_108), inference(avatar_split_clause, [], [f3673, f897, f823, f571])).
fof(f897, plain, (spl15_108 <=> (e1 = op(op(e0, op(e0, e0)), e0))), introduced(avatar_definition, [new_symbols(naming, [spl15_108])])).
fof(f3673, plain, ((op(e0, e0) = e1) | (~ spl15_92 | ~ spl15_108)), inference(backward_demodulation, [], [f898, f825])).
fof(f898, plain, ((e1 = op(op(e0, op(e0, e0)), e0)) | ~ spl15_108), inference(avatar_component_clause, [], [f897])).
fof(f3672, plain, (~ spl15_18 | ~ spl15_97 | spl15_105), inference(avatar_split_clause, [], [f3669, f882, f844, f384])).
fof(f844, plain, (spl15_97 <=> (e2 = op(e3, op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl15_97])])).
fof(f882, plain, (spl15_105 <=> (e1 = op(op(e3, op(e3, e3)), e3))), introduced(avatar_definition, [new_symbols(naming, [spl15_105])])).
fof(f3669, plain, (~ (e1 = op(e2, e3)) | (~ spl15_97 | spl15_105)), inference(backward_demodulation, [], [f884, f845])).
fof(f845, plain, ((e2 = op(e3, op(e3, e3))) | ~ spl15_97), inference(avatar_component_clause, [], [f844])).
fof(f884, plain, (~ (e1 = op(op(e3, op(e3, e3)), e3)) | spl15_105), inference(avatar_component_clause, [], [f882])).
fof(f3667, plain, (~ spl15_5 | spl15_98 | ~ spl15_99), inference(avatar_split_clause, [], [f3666, f853, f849, f329])).
fof(f849, plain, (spl15_98 <=> (e0 = op(op(e2, op(e2, e2)), e2))), introduced(avatar_definition, [new_symbols(naming, [spl15_98])])).
fof(f853, plain, (spl15_99 <=> (e3 = op(e2, op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl15_99])])).
fof(f3666, plain, (~ (e0 = op(e3, e2)) | (spl15_98 | ~ spl15_99)), inference(forward_demodulation, [], [f851, f854])).
fof(f854, plain, ((e3 = op(e2, op(e2, e2))) | ~ spl15_99), inference(avatar_component_clause, [], [f853])).
fof(f851, plain, (~ (e0 = op(op(e2, op(e2, e2)), e2)) | spl15_98), inference(avatar_component_clause, [], [f849])).
fof(f3646, plain, (~ spl15_6 | ~ spl15_14), inference(avatar_split_clause, [], [f3645, f367, f333])).
fof(f333, plain, (spl15_6 <=> (e1 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl15_6])])).
fof(f3645, plain, (~ (e1 = op(e3, e2)) | ~ spl15_14), inference(backward_demodulation, [], [f176, f369])).
fof(f369, plain, ((e1 = op(e3, e0)) | ~ spl15_14), inference(avatar_component_clause, [], [f367])).
fof(f3638, plain, (~ spl15_6 | ~ spl15_38), inference(avatar_split_clause, [], [f3631, f469, f333])).
fof(f469, plain, (spl15_38 <=> (e1 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl15_38])])).
fof(f3631, plain, (~ (e1 = op(e3, e2)) | ~ spl15_38), inference(backward_demodulation, [], [f149, f471])).
fof(f471, plain, ((e1 = op(e1, e2)) | ~ spl15_38), inference(avatar_component_clause, [], [f469])).
fof(f3628, plain, (~ spl15_13 | ~ spl15_45), inference(avatar_split_clause, [], [f3623, f499, f363])).
fof(f3623, plain, (~ (e0 = op(e3, e0)) | ~ spl15_45), inference(backward_demodulation, [], [f137, f501])).
fof(f137, plain, ~ (op(e1, e0) = op(e3, e0)), inference(cnf_transformation, [], [f3])).
fof(f3619, plain, (~ spl15_3 | spl15_8 | ~ spl15_83), inference(avatar_contradiction_clause, [], [f3618])).
fof(f3618, plain, ($false | (~ spl15_3 | spl15_8 | ~ spl15_83)), inference(subsumption_resolution, [], [f3617, f342])).
fof(f342, plain, (~ (e3 = op(e3, e2)) | spl15_8), inference(avatar_component_clause, [], [f341])).
fof(f3617, plain, ((e3 = op(e3, e2)) | (~ spl15_3 | ~ spl15_83)), inference(forward_demodulation, [], [f786, f322])).
fof(f322, plain, ((e2 = op(e3, e3)) | ~ spl15_3), inference(avatar_component_clause, [], [f320])).
fof(f3616, plain, (~ spl15_25 | spl15_31 | ~ spl15_90), inference(avatar_contradiction_clause, [], [f3615])).
fof(f3615, plain, ($false | (~ spl15_25 | spl15_31 | ~ spl15_90)), inference(subsumption_resolution, [], [f3614, f440])).
fof(f440, plain, (~ (e2 = op(e2, e0)) | spl15_31), inference(avatar_component_clause, [], [f439])).
fof(f439, plain, (spl15_31 <=> (e2 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl15_31])])).
fof(f3614, plain, ((e2 = op(e2, e0)) | (~ spl15_25 | ~ spl15_90)), inference(forward_demodulation, [], [f816, f416])).
fof(f416, plain, ((e0 = op(e2, e1)) | ~ spl15_25), inference(avatar_component_clause, [], [f414])).
fof(f3609, plain, (~ spl15_37 | ~ spl15_43 | spl15_115), inference(avatar_split_clause, [], [f3608, f932, f490, f465])).
fof(f490, plain, (spl15_43 <=> (e2 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl15_43])])).
fof(f932, plain, (spl15_115 <=> (e0 = op(e1, op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl15_115])])).
fof(f3608, plain, (~ (e0 = op(e1, e2)) | (~ spl15_43 | spl15_115)), inference(forward_demodulation, [], [f934, f492])).
fof(f492, plain, ((e2 = op(e1, e1)) | ~ spl15_43), inference(avatar_component_clause, [], [f490])).
fof(f934, plain, (~ (e0 = op(e1, op(e1, e1))) | spl15_115), inference(avatar_component_clause, [], [f932])).
fof(f3603, plain, (~ spl15_38 | ~ spl15_43 | spl15_89), inference(avatar_split_clause, [], [f3602, f810, f490, f469])).
fof(f810, plain, (spl15_89 <=> (e1 = op(e1, op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl15_89])])).
fof(f3602, plain, (~ (e1 = op(e1, e2)) | (~ spl15_43 | spl15_89)), inference(forward_demodulation, [], [f811, f492])).
fof(f811, plain, (~ (e1 = op(e1, op(e1, e1))) | spl15_89), inference(avatar_component_clause, [], [f810])).
fof(f3601, plain, (~ spl15_47 | ~ spl15_43), inference(avatar_split_clause, [], [f3600, f490, f507])).
fof(f507, plain, (spl15_47 <=> (e2 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl15_47])])).
fof(f3600, plain, (~ (e2 = op(e1, e0)) | ~ spl15_43), inference(forward_demodulation, [], [f163, f492])).
fof(f163, plain, ~ (op(e1, e0) = op(e1, e1)), inference(cnf_transformation, [], [f3])).
fof(f3599, plain, (~ spl15_47 | ~ spl15_63), inference(avatar_split_clause, [], [f3598, f575, f507])).
fof(f3598, plain, (~ (e2 = op(e1, e0)) | ~ spl15_63), inference(forward_demodulation, [], [f133, f577])).
fof(f3597, plain, (~ spl15_39 | ~ spl15_23), inference(avatar_split_clause, [], [f3596, f405, f473])).
fof(f473, plain, (spl15_39 <=> (e2 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl15_39])])).
fof(f3596, plain, (~ (e2 = op(e1, e2)) | ~ spl15_23), inference(forward_demodulation, [], [f148, f407])).
fof(f407, plain, ((e2 = op(e2, e2)) | ~ spl15_23), inference(avatar_component_clause, [], [f405])).
fof(f148, plain, ~ (op(e1, e2) = op(e2, e2)), inference(cnf_transformation, [], [f3])).
fof(f3595, plain, (~ spl15_14 | spl15_108 | ~ spl15_109), inference(avatar_split_clause, [], [f3506, f901, f897, f367])).
fof(f901, plain, (spl15_109 <=> (e3 = op(e0, op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl15_109])])).
fof(f3506, plain, (~ (e1 = op(e3, e0)) | (spl15_108 | ~ spl15_109)), inference(forward_demodulation, [], [f899, f902])).
fof(f902, plain, ((e3 = op(e0, op(e0, e0))) | ~ spl15_109), inference(avatar_component_clause, [], [f901])).
fof(f899, plain, (~ (e1 = op(op(e0, op(e0, e0)), e0)) | spl15_108), inference(avatar_component_clause, [], [f897])).
fof(f3592, plain, (~ spl15_7 | ~ spl15_23), inference(avatar_split_clause, [], [f3591, f405, f337])).
fof(f337, plain, (spl15_7 <=> (e2 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl15_7])])).
fof(f3591, plain, (~ (e2 = op(e3, e2)) | ~ spl15_23), inference(forward_demodulation, [], [f150, f407])).
fof(f3589, plain, (~ spl15_5 | ~ spl15_3 | spl15_107), inference(avatar_split_clause, [], [f3588, f892, f320, f329])).
fof(f892, plain, (spl15_107 <=> (e0 = op(e3, op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl15_107])])).
fof(f3588, plain, (~ (e0 = op(e3, e2)) | (~ spl15_3 | spl15_107)), inference(backward_demodulation, [], [f894, f322])).
fof(f894, plain, (~ (e0 = op(e3, op(e3, e3))) | spl15_107), inference(avatar_component_clause, [], [f892])).
fof(f3584, plain, (~ spl15_1 | ~ spl15_13), inference(avatar_split_clause, [], [f3579, f363, f312])).
fof(f3579, plain, (~ (e0 = op(e3, e3)) | ~ spl15_13), inference(backward_demodulation, [], [f177, f365])).
fof(f3578, plain, (~ spl15_18 | ~ spl15_25 | ~ spl15_82), inference(avatar_contradiction_clause, [], [f3577])).
fof(f3577, plain, ($false | (~ spl15_18 | ~ spl15_25 | ~ spl15_82)), inference(subsumption_resolution, [], [f3576, f182])).
fof(f182, plain, ~ (e0 = e2), inference(cnf_transformation, [], [f4])).
fof(f3576, plain, ((e0 = e2) | (~ spl15_18 | ~ spl15_25 | ~ spl15_82)), inference(forward_demodulation, [], [f3575, f416])).
fof(f3575, plain, ((e2 = op(e2, e1)) | (~ spl15_18 | ~ spl15_82)), inference(backward_demodulation, [], [f782, f386])).
fof(f386, plain, ((e1 = op(e2, e3)) | ~ spl15_18), inference(avatar_component_clause, [], [f384])).
fof(f782, plain, ((e2 = op(e2, op(e2, e3))) | ~ spl15_82), inference(avatar_component_clause, [], [f780])).
fof(f780, plain, (spl15_82 <=> (e2 = op(e2, op(e2, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl15_82])])).
fof(f3570, plain, (~ spl15_19 | ~ spl15_23), inference(avatar_split_clause, [], [f3565, f405, f388])).
fof(f3565, plain, (~ (e2 = op(e2, e3)) | ~ spl15_23), inference(backward_demodulation, [], [f174, f407])).
fof(f3564, plain, (~ spl15_21 | ~ spl15_25), inference(avatar_split_clause, [], [f3558, f414, f397])).
fof(f397, plain, (spl15_21 <=> (e0 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl15_21])])).
fof(f3558, plain, (~ (e0 = op(e2, e2)) | ~ spl15_25), inference(backward_demodulation, [], [f172, f416])).
fof(f172, plain, ~ (op(e2, e1) = op(e2, e2)), inference(cnf_transformation, [], [f3])).
fof(f3557, plain, (~ spl15_21 | ~ spl15_37), inference(avatar_split_clause, [], [f3553, f465, f397])).
fof(f3553, plain, (~ (e0 = op(e2, e2)) | ~ spl15_37), inference(backward_demodulation, [], [f148, f467])).
fof(f3556, plain, (~ spl15_37 | ~ spl15_46 | spl15_85), inference(avatar_contradiction_clause, [], [f3555])).
fof(f3555, plain, ($false | (~ spl15_37 | ~ spl15_46 | spl15_85)), inference(subsumption_resolution, [], [f3551, f505])).
fof(f3551, plain, (~ (e1 = op(e1, e0)) | (~ spl15_37 | spl15_85)), inference(backward_demodulation, [], [f794, f467])).
fof(f794, plain, (~ (e1 = op(e1, op(e1, e2))) | spl15_85), inference(avatar_component_clause, [], [f793])).
fof(f3548, plain, (~ spl15_39 | ~ spl15_43), inference(avatar_split_clause, [], [f3545, f490, f473])).
fof(f3545, plain, (~ (e2 = op(e1, e2)) | ~ spl15_43), inference(backward_demodulation, [], [f166, f492])).
fof(f166, plain, ~ (op(e1, e1) = op(e1, e2)), inference(cnf_transformation, [], [f3])).
fof(f3543, plain, (~ spl15_1 | ~ spl15_49), inference(avatar_split_clause, [], [f3538, f516, f312])).
fof(f3538, plain, (~ (e0 = op(e3, e3)) | ~ spl15_49), inference(backward_demodulation, [], [f153, f518])).
fof(f518, plain, ((e0 = op(e0, e3)) | ~ spl15_49), inference(avatar_component_clause, [], [f516])).
fof(f153, plain, ~ (op(e0, e3) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f3536, plain, (~ spl15_50 | ~ spl15_58), inference(avatar_split_clause, [], [f3531, f554, f520])).
fof(f3531, plain, (~ (e1 = op(e0, e3)) | ~ spl15_58), inference(backward_demodulation, [], [f161, f556])).
fof(f556, plain, ((e1 = op(e0, e1)) | ~ spl15_58), inference(avatar_component_clause, [], [f554])).
fof(f3535, plain, (~ spl15_26 | ~ spl15_58), inference(avatar_split_clause, [], [f3530, f554, f418])).
fof(f418, plain, (spl15_26 <=> (e1 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl15_26])])).
fof(f3530, plain, (~ (e1 = op(e2, e1)) | ~ spl15_58), inference(backward_demodulation, [], [f140, f556])).
fof(f3528, plain, (~ spl15_59 | ~ spl15_63), inference(avatar_split_clause, [], [f3521, f575, f558])).
fof(f558, plain, (spl15_59 <=> (e2 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl15_59])])).
fof(f3521, plain, (~ (e2 = op(e0, e1)) | ~ spl15_63), inference(backward_demodulation, [], [f157, f577])).
fof(f3519, plain, (spl15_49 | ~ spl15_56 | ~ spl15_84), inference(avatar_split_clause, [], [f3518, f789, f545, f516])).
fof(f545, plain, (spl15_56 <=> (e3 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl15_56])])).
fof(f789, plain, (spl15_84 <=> (e0 = op(e0, op(e0, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl15_84])])).
fof(f3518, plain, ((e0 = op(e0, e3)) | (~ spl15_56 | ~ spl15_84)), inference(forward_demodulation, [], [f791, f547])).
fof(f547, plain, ((e3 = op(e0, e2)) | ~ spl15_56), inference(avatar_component_clause, [], [f545])).
fof(f791, plain, ((e0 = op(e0, op(e0, e2))) | ~ spl15_84), inference(avatar_component_clause, [], [f789])).
fof(f3513, plain, (spl15_4 | ~ spl15_12 | ~ spl15_91), inference(avatar_contradiction_clause, [], [f3512])).
fof(f3512, plain, ($false | (spl15_4 | ~ spl15_12 | ~ spl15_91)), inference(subsumption_resolution, [], [f3511, f325])).
fof(f325, plain, (~ (e3 = op(e3, e3)) | spl15_4), inference(avatar_component_clause, [], [f324])).
fof(f3511, plain, ((e3 = op(e3, e3)) | (~ spl15_12 | ~ spl15_91)), inference(forward_demodulation, [], [f820, f360])).
fof(f360, plain, ((e3 = op(e3, e1)) | ~ spl15_12), inference(avatar_component_clause, [], [f358])).
fof(f820, plain, ((e3 = op(e3, op(e3, e1))) | ~ spl15_91), inference(avatar_component_clause, [], [f818])).
fof(f818, plain, (spl15_91 <=> (e3 = op(e3, op(e3, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl15_91])])).
fof(f3502, plain, (~ spl15_59 | spl15_113 | ~ spl15_115), inference(avatar_split_clause, [], [f3500, f932, f921, f558])).
fof(f921, plain, (spl15_113 <=> (e2 = op(op(e1, op(e1, e1)), e1))), introduced(avatar_definition, [new_symbols(naming, [spl15_113])])).
fof(f3500, plain, (~ (e2 = op(e0, e1)) | (spl15_113 | ~ spl15_115)), inference(backward_demodulation, [], [f923, f933])).
fof(f933, plain, ((e0 = op(e1, op(e1, e1))) | ~ spl15_115), inference(avatar_component_clause, [], [f932])).
fof(f923, plain, (~ (e2 = op(op(e1, op(e1, e1)), e1)) | spl15_113), inference(avatar_component_clause, [], [f921])).
fof(f3493, plain, (~ spl15_42 | ~ spl15_46), inference(avatar_split_clause, [], [f3492, f503, f486])).
fof(f3492, plain, (~ (e1 = op(e1, e1)) | ~ spl15_46), inference(forward_demodulation, [], [f163, f505])).
fof(f3487, plain, (~ spl15_22 | ~ spl15_6), inference(avatar_split_clause, [], [f3486, f333, f401])).
fof(f3486, plain, (~ (e1 = op(e2, e2)) | ~ spl15_6), inference(forward_demodulation, [], [f150, f335])).
fof(f335, plain, ((e1 = op(e3, e2)) | ~ spl15_6), inference(avatar_component_clause, [], [f333])).
fof(f3485, plain, (~ spl15_20 | ~ spl15_32), inference(avatar_split_clause, [], [f3484, f443, f392])).
fof(f392, plain, (spl15_20 <=> (e3 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl15_20])])).
fof(f3484, plain, (~ (e3 = op(e2, e3)) | ~ spl15_32), inference(forward_demodulation, [], [f171, f445])).
fof(f171, plain, ~ (op(e2, e0) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f3483, plain, (spl15_19 | ~ spl15_32 | ~ spl15_94), inference(avatar_split_clause, [], [f3415, f831, f443, f388])).
fof(f3415, plain, ((e2 = op(e2, e3)) | (~ spl15_32 | ~ spl15_94)), inference(forward_demodulation, [], [f833, f445])).
fof(f3475, plain, (~ spl15_2 | ~ spl15_6), inference(avatar_split_clause, [], [f3471, f333, f316])).
fof(f316, plain, (spl15_2 <=> (e1 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl15_2])])).
fof(f3471, plain, (~ (e1 = op(e3, e3)) | ~ spl15_6), inference(backward_demodulation, [], [f180, f335])).
fof(f3470, plain, (~ spl15_7 | ~ spl15_15), inference(avatar_split_clause, [], [f3466, f371, f337])).
fof(f3466, plain, (~ (e2 = op(e3, e2)) | ~ spl15_15), inference(backward_demodulation, [], [f176, f373])).
fof(f3465, plain, (~ spl15_7 | ~ spl15_39), inference(avatar_split_clause, [], [f3463, f473, f337])).
fof(f3463, plain, (~ (e2 = op(e3, e2)) | ~ spl15_39), inference(backward_demodulation, [], [f149, f475])).
fof(f475, plain, ((e2 = op(e1, e2)) | ~ spl15_39), inference(avatar_component_clause, [], [f473])).
fof(f3464, plain, (~ spl15_38 | ~ spl15_39 | spl15_85), inference(avatar_split_clause, [], [f3461, f793, f473, f469])).
fof(f3461, plain, (~ (e1 = op(e1, e2)) | (~ spl15_39 | spl15_85)), inference(backward_demodulation, [], [f794, f475])).
fof(f3455, plain, (~ spl15_38 | ~ spl15_46), inference(avatar_split_clause, [], [f3448, f503, f469])).
fof(f3448, plain, (~ (e1 = op(e1, e2)) | ~ spl15_46), inference(backward_demodulation, [], [f164, f505])).
fof(f3446, plain, (~ spl15_2 | ~ spl15_50), inference(avatar_split_clause, [], [f3441, f520, f316])).
fof(f3441, plain, (~ (e1 = op(e3, e3)) | ~ spl15_50), inference(backward_demodulation, [], [f153, f522])).
fof(f522, plain, ((e1 = op(e0, e3)) | ~ spl15_50), inference(avatar_component_clause, [], [f520])).
fof(f3445, plain, (~ spl15_50 | spl15_57 | ~ spl15_80), inference(avatar_contradiction_clause, [], [f3444])).
fof(f3444, plain, ($false | (~ spl15_50 | spl15_57 | ~ spl15_80)), inference(subsumption_resolution, [], [f3440, f551])).
fof(f551, plain, (~ (e0 = op(e0, e1)) | spl15_57), inference(avatar_component_clause, [], [f550])).
fof(f3440, plain, ((e0 = op(e0, e1)) | (~ spl15_50 | ~ spl15_80)), inference(backward_demodulation, [], [f774, f522])).
fof(f774, plain, ((e0 = op(e0, op(e0, e3))) | ~ spl15_80), inference(avatar_component_clause, [], [f772])).
fof(f3433, plain, (~ spl15_61 | spl15_92), inference(avatar_contradiction_clause, [], [f3432])).
fof(f3432, plain, ($false | (~ spl15_61 | spl15_92)), inference(subsumption_resolution, [], [f3425, f569])).
fof(f569, plain, ((e0 = op(e0, e0)) | ~ spl15_61), inference(avatar_component_clause, [], [f567])).
fof(f567, plain, (spl15_61 <=> (e0 = op(e0, e0))), introduced(avatar_definition, [new_symbols(naming, [spl15_61])])).
fof(f3425, plain, (~ (e0 = op(e0, e0)) | (~ spl15_61 | spl15_92)), inference(backward_demodulation, [], [f824, f569])).
fof(f824, plain, (~ (e0 = op(e0, op(e0, e0))) | spl15_92), inference(avatar_component_clause, [], [f823])).
fof(f3431, plain, (~ spl15_49 | ~ spl15_61), inference(avatar_split_clause, [], [f3421, f567, f516])).
fof(f3421, plain, (~ (e0 = op(e0, e3)) | ~ spl15_61), inference(backward_demodulation, [], [f159, f569])).
fof(f3430, plain, (~ spl15_13 | ~ spl15_61), inference(avatar_split_clause, [], [f3420, f567, f363])).
fof(f3420, plain, (~ (e0 = op(e3, e0)) | ~ spl15_61), inference(backward_demodulation, [], [f135, f569])).
fof(f3409, plain, (~ spl15_21 | ~ spl15_32 | spl15_99), inference(avatar_contradiction_clause, [], [f3408])).
fof(f3408, plain, ($false | (~ spl15_21 | ~ spl15_32 | spl15_99)), inference(subsumption_resolution, [], [f3407, f445])).
fof(f3407, plain, (~ (e3 = op(e2, e0)) | (~ spl15_21 | spl15_99)), inference(forward_demodulation, [], [f855, f399])).
fof(f399, plain, ((e0 = op(e2, e2)) | ~ spl15_21), inference(avatar_component_clause, [], [f397])).
fof(f855, plain, (~ (e3 = op(e2, op(e2, e2))) | spl15_99), inference(avatar_component_clause, [], [f853])).
fof(f3406, plain, (~ spl15_47 | ~ spl15_41 | spl15_104), inference(avatar_split_clause, [], [f3405, f877, f482, f507])).
fof(f482, plain, (spl15_41 <=> (e0 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl15_41])])).
fof(f877, plain, (spl15_104 <=> (e2 = op(e1, op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl15_104])])).
fof(f3405, plain, (~ (e2 = op(e1, e0)) | (~ spl15_41 | spl15_104)), inference(forward_demodulation, [], [f879, f484])).
fof(f484, plain, ((e0 = op(e1, e1)) | ~ spl15_41), inference(avatar_component_clause, [], [f482])).
fof(f879, plain, (~ (e2 = op(e1, op(e1, e1))) | spl15_104), inference(avatar_component_clause, [], [f877])).
fof(f3395, plain, (spl15_46 | ~ spl15_41 | ~ spl15_89), inference(avatar_split_clause, [], [f3334, f810, f482, f503])).
fof(f3334, plain, ((e1 = op(e1, e0)) | (~ spl15_41 | ~ spl15_89)), inference(backward_demodulation, [], [f812, f484])).
fof(f812, plain, ((e1 = op(e1, op(e1, e1))) | ~ spl15_89), inference(avatar_component_clause, [], [f810])).
fof(f3392, plain, (~ spl15_8 | ~ spl15_12), inference(avatar_split_clause, [], [f3391, f358, f341])).
fof(f3391, plain, (~ (e3 = op(e3, e2)) | ~ spl15_12), inference(forward_demodulation, [], [f178, f360])).
fof(f178, plain, ~ (op(e3, e1) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f3361, plain, (~ spl15_22 | ~ spl15_26), inference(avatar_split_clause, [], [f3357, f418, f401])).
fof(f3357, plain, (~ (e1 = op(e2, e2)) | ~ spl15_26), inference(backward_demodulation, [], [f172, f420])).
fof(f420, plain, ((e1 = op(e2, e1)) | ~ spl15_26), inference(avatar_component_clause, [], [f418])).
fof(f3360, plain, (~ spl15_10 | ~ spl15_26), inference(avatar_split_clause, [], [f3356, f418, f350])).
fof(f350, plain, (spl15_10 <=> (e1 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl15_10])])).
fof(f3356, plain, (~ (e1 = op(e3, e1)) | ~ spl15_26), inference(backward_demodulation, [], [f144, f420])).
fof(f3353, plain, (~ spl15_19 | ~ spl15_32 | spl15_94), inference(avatar_contradiction_clause, [], [f3352])).
fof(f3352, plain, ($false | (~ spl15_19 | ~ spl15_32 | spl15_94)), inference(subsumption_resolution, [], [f3349, f390])).
fof(f3349, plain, (~ (e2 = op(e2, e3)) | (~ spl15_32 | spl15_94)), inference(backward_demodulation, [], [f832, f445])).
fof(f832, plain, (~ (e2 = op(e2, op(e2, e0))) | spl15_94), inference(avatar_component_clause, [], [f831])).
fof(f3351, plain, (~ spl15_28 | ~ spl15_32), inference(avatar_split_clause, [], [f3347, f443, f426])).
fof(f3347, plain, (~ (e3 = op(e2, e1)) | ~ spl15_32), inference(backward_demodulation, [], [f169, f445])).
fof(f3350, plain, (~ spl15_16 | ~ spl15_32), inference(avatar_split_clause, [], [f3346, f443, f375])).
fof(f375, plain, (spl15_16 <=> (e3 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl15_16])])).
fof(f3346, plain, (~ (e3 = op(e3, e0)) | ~ spl15_32), inference(backward_demodulation, [], [f138, f445])).
fof(f138, plain, ~ (op(e2, e0) = op(e3, e0)), inference(cnf_transformation, [], [f3])).
fof(f3345, plain, (~ spl15_22 | ~ spl15_38), inference(avatar_split_clause, [], [f3341, f469, f401])).
fof(f3341, plain, (~ (e1 = op(e2, e2)) | ~ spl15_38), inference(backward_demodulation, [], [f148, f471])).
fof(f3337, plain, (~ spl15_37 | ~ spl15_41), inference(avatar_split_clause, [], [f3330, f482, f465])).
fof(f3330, plain, (~ (e0 = op(e1, e2)) | ~ spl15_41), inference(backward_demodulation, [], [f166, f484])).
fof(f3336, plain, (~ spl15_25 | ~ spl15_41), inference(avatar_split_clause, [], [f3328, f482, f414])).
fof(f3328, plain, (~ (e0 = op(e2, e1)) | ~ spl15_41), inference(backward_demodulation, [], [f142, f484])).
fof(f3327, plain, (spl15_38 | ~ spl15_47 | ~ spl15_93), inference(avatar_split_clause, [], [f3324, f827, f507, f469])).
fof(f3324, plain, ((e1 = op(e1, e2)) | (~ spl15_47 | ~ spl15_93)), inference(backward_demodulation, [], [f829, f509])).
fof(f509, plain, ((e2 = op(e1, e0)) | ~ spl15_47), inference(avatar_component_clause, [], [f507])).
fof(f829, plain, ((e1 = op(e1, op(e1, e0))) | ~ spl15_93), inference(avatar_component_clause, [], [f827])).
fof(f3325, plain, (~ spl15_31 | ~ spl15_47), inference(avatar_split_clause, [], [f3319, f507, f439])).
fof(f3319, plain, (~ (e2 = op(e2, e0)) | ~ spl15_47), inference(backward_demodulation, [], [f136, f509])).
fof(f3312, plain, (~ spl15_50 | ~ spl15_62), inference(avatar_split_clause, [], [f3306, f571, f520])).
fof(f3306, plain, (~ (e1 = op(e0, e3)) | ~ spl15_62), inference(backward_demodulation, [], [f159, f573])).
fof(f3311, plain, (~ spl15_46 | ~ spl15_62), inference(avatar_split_clause, [], [f3303, f571, f503])).
fof(f3303, plain, (~ (e1 = op(e1, e0)) | ~ spl15_62), inference(backward_demodulation, [], [f133, f573])).
fof(f3300, plain, (~ spl15_89 | ~ spl15_104), inference(avatar_contradiction_clause, [], [f3299])).
fof(f3299, plain, ($false | (~ spl15_89 | ~ spl15_104)), inference(subsumption_resolution, [], [f3298, f184])).
fof(f184, plain, ~ (e1 = e2), inference(cnf_transformation, [], [f4])).
fof(f3298, plain, ((e1 = e2) | (~ spl15_89 | ~ spl15_104)), inference(backward_demodulation, [], [f878, f812])).
fof(f878, plain, ((e2 = op(e1, op(e1, e1))) | ~ spl15_104), inference(avatar_component_clause, [], [f877])).
fof(f3295, plain, (~ spl15_25 | spl15_101 | ~ spl15_104), inference(avatar_split_clause, [], [f3293, f877, f863, f414])).
fof(f863, plain, (spl15_101 <=> (e0 = op(op(e1, op(e1, e1)), e1))), introduced(avatar_definition, [new_symbols(naming, [spl15_101])])).
fof(f3293, plain, (~ (e0 = op(e2, e1)) | (spl15_101 | ~ spl15_104)), inference(backward_demodulation, [], [f865, f878])).
fof(f865, plain, (~ (e0 = op(op(e1, op(e1, e1)), e1)) | spl15_101), inference(avatar_component_clause, [], [f863])).
fof(f3294, plain, (~ spl15_28 | ~ spl15_104 | spl15_118), inference(avatar_split_clause, [], [f3292, f947, f877, f426])).
fof(f947, plain, (spl15_118 <=> (e3 = op(op(e1, op(e1, e1)), e1))), introduced(avatar_definition, [new_symbols(naming, [spl15_118])])).
fof(f3292, plain, (~ (e3 = op(e2, e1)) | (~ spl15_104 | spl15_118)), inference(backward_demodulation, [], [f949, f878])).
fof(f949, plain, (~ (e3 = op(op(e1, op(e1, e1)), e1)) | spl15_118), inference(avatar_component_clause, [], [f947])).
fof(f3289, plain, (~ spl15_32 | ~ spl15_111 | spl15_119), inference(avatar_split_clause, [], [f3287, f953, f911, f443])).
fof(f911, plain, (spl15_111 <=> (e2 = op(e0, op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl15_111])])).
fof(f953, plain, (spl15_119 <=> (e3 = op(op(e0, op(e0, e0)), e0))), introduced(avatar_definition, [new_symbols(naming, [spl15_119])])).
fof(f3287, plain, (~ (e3 = op(e2, e0)) | (~ spl15_111 | spl15_119)), inference(backward_demodulation, [], [f955, f912])).
fof(f912, plain, ((e2 = op(e0, op(e0, e0))) | ~ spl15_111), inference(avatar_component_clause, [], [f911])).
fof(f955, plain, (~ (e3 = op(op(e0, op(e0, e0)), e0)) | spl15_119), inference(avatar_component_clause, [], [f953])).
fof(f3274, plain, (~ spl15_40 | ~ spl15_56), inference(avatar_split_clause, [], [f3273, f545, f477])).
fof(f477, plain, (spl15_40 <=> (e3 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl15_40])])).
fof(f3273, plain, (~ (e3 = op(e1, e2)) | ~ spl15_56), inference(forward_demodulation, [], [f145, f547])).
fof(f145, plain, ~ (op(e0, e2) = op(e1, e2)), inference(cnf_transformation, [], [f3])).
fof(f3272, plain, (~ spl15_31 | ~ spl15_19), inference(avatar_split_clause, [], [f3271, f388, f439])).
fof(f3271, plain, (~ (e2 = op(e2, e0)) | ~ spl15_19), inference(forward_demodulation, [], [f171, f390])).
fof(f3269, plain, (~ spl15_27 | ~ spl15_59), inference(avatar_split_clause, [], [f3268, f558, f422])).
fof(f3268, plain, (~ (e2 = op(e2, e1)) | ~ spl15_59), inference(forward_demodulation, [], [f140, f560])).
fof(f560, plain, ((e2 = op(e0, e1)) | ~ spl15_59), inference(avatar_component_clause, [], [f558])).
fof(f3265, plain, (~ spl15_24 | ~ spl15_56), inference(avatar_split_clause, [], [f3264, f545, f409])).
fof(f3264, plain, (~ (e3 = op(e2, e2)) | ~ spl15_56), inference(forward_demodulation, [], [f146, f547])).
fof(f3250, plain, (~ spl15_2 | ~ spl15_10), inference(avatar_split_clause, [], [f3246, f350, f316])).
fof(f3246, plain, (~ (e1 = op(e3, e3)) | ~ spl15_10), inference(backward_demodulation, [], [f179, f352])).
fof(f352, plain, ((e1 = op(e3, e1)) | ~ spl15_10), inference(avatar_component_clause, [], [f350])).
fof(f179, plain, ~ (op(e3, e1) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f3244, plain, (~ spl15_8 | ~ spl15_16), inference(avatar_split_clause, [], [f3241, f375, f341])).
fof(f3241, plain, (~ (e3 = op(e3, e2)) | ~ spl15_16), inference(backward_demodulation, [], [f176, f377])).
fof(f377, plain, ((e3 = op(e3, e0)) | ~ spl15_16), inference(avatar_component_clause, [], [f375])).
fof(f3232, plain, (~ spl15_20 | ~ spl15_36), inference(avatar_split_clause, [], [f3229, f460, f392])).
fof(f3229, plain, (~ (e3 = op(e2, e3)) | ~ spl15_36), inference(backward_demodulation, [], [f154, f462])).
fof(f3228, plain, (~ spl15_35 | ~ spl15_43), inference(avatar_split_clause, [], [f3224, f490, f456])).
fof(f456, plain, (spl15_35 <=> (e2 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl15_35])])).
fof(f3224, plain, (~ (e2 = op(e1, e3)) | ~ spl15_43), inference(backward_demodulation, [], [f167, f492])).
fof(f167, plain, ~ (op(e1, e1) = op(e1, e3)), inference(cnf_transformation, [], [f3])).
fof(f3227, plain, (~ spl15_11 | ~ spl15_43), inference(avatar_split_clause, [], [f3223, f490, f354])).
fof(f354, plain, (spl15_11 <=> (e2 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl15_11])])).
fof(f3223, plain, (~ (e2 = op(e3, e1)) | ~ spl15_43), inference(backward_demodulation, [], [f143, f492])).
fof(f143, plain, ~ (op(e1, e1) = op(e3, e1)), inference(cnf_transformation, [], [f3])).
fof(f3215, plain, (~ spl15_8 | ~ spl15_56), inference(avatar_split_clause, [], [f3212, f545, f341])).
fof(f3212, plain, (~ (e3 = op(e3, e2)) | ~ spl15_56), inference(backward_demodulation, [], [f147, f547])).
fof(f3211, plain, (~ spl15_55 | ~ spl15_59), inference(avatar_split_clause, [], [f3205, f558, f541])).
fof(f3205, plain, (~ (e2 = op(e0, e2)) | ~ spl15_59), inference(backward_demodulation, [], [f160, f560])).
fof(f160, plain, ~ (op(e0, e1) = op(e0, e2)), inference(cnf_transformation, [], [f3])).
fof(f3210, plain, (~ spl15_11 | ~ spl15_59), inference(avatar_split_clause, [], [f3204, f558, f354])).
fof(f3204, plain, (~ (e2 = op(e3, e1)) | ~ spl15_59), inference(backward_demodulation, [], [f141, f560])).
fof(f141, plain, ~ (op(e0, e1) = op(e3, e1)), inference(cnf_transformation, [], [f3])).
fof(f3209, plain, (~ spl15_43 | ~ spl15_59), inference(avatar_split_clause, [], [f3203, f558, f490])).
fof(f3203, plain, (~ (e2 = op(e1, e1)) | ~ spl15_59), inference(backward_demodulation, [], [f139, f560])).
fof(f3171, plain, (~ spl15_56 | ~ spl15_110 | spl15_117), inference(avatar_split_clause, [], [f3050, f942, f906, f545])).
fof(f942, plain, (spl15_117 <=> (e3 = op(op(e2, op(e2, e2)), e2))), introduced(avatar_definition, [new_symbols(naming, [spl15_117])])).
fof(f3050, plain, (~ (e3 = op(e0, e2)) | (~ spl15_110 | spl15_117)), inference(backward_demodulation, [], [f944, f907])).
fof(f907, plain, ((e0 = op(e2, op(e2, e2))) | ~ spl15_110), inference(avatar_component_clause, [], [f906])).
fof(f944, plain, (~ (e3 = op(op(e2, op(e2, e2)), e2)) | spl15_117), inference(avatar_component_clause, [], [f942])).
fof(f3157, plain, (spl15_8 | ~ spl15_11 | ~ spl15_91), inference(avatar_split_clause, [], [f3155, f818, f354, f341])).
fof(f3155, plain, ((e3 = op(e3, e2)) | (~ spl15_11 | ~ spl15_91)), inference(backward_demodulation, [], [f820, f356])).
fof(f356, plain, ((e2 = op(e3, e1)) | ~ spl15_11), inference(avatar_component_clause, [], [f354])).
fof(f3156, plain, (~ spl15_3 | ~ spl15_11), inference(avatar_split_clause, [], [f3153, f354, f320])).
fof(f3153, plain, (~ (e2 = op(e3, e3)) | ~ spl15_11), inference(backward_demodulation, [], [f179, f356])).
fof(f3150, plain, (~ spl15_9 | ~ spl15_13), inference(avatar_split_clause, [], [f3143, f363, f346])).
fof(f346, plain, (spl15_9 <=> (e0 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl15_9])])).
fof(f3143, plain, (~ (e0 = op(e3, e1)) | ~ spl15_13), inference(backward_demodulation, [], [f175, f365])).
fof(f175, plain, ~ (op(e3, e0) = op(e3, e1)), inference(cnf_transformation, [], [f3])).
fof(f3131, plain, (~ spl15_9 | ~ spl15_25), inference(avatar_split_clause, [], [f3124, f414, f346])).
fof(f3124, plain, (~ (e0 = op(e3, e1)) | ~ spl15_25), inference(backward_demodulation, [], [f144, f416])).
fof(f3122, plain, (~ spl15_23 | ~ spl15_31), inference(avatar_split_clause, [], [f3118, f439, f405])).
fof(f3118, plain, (~ (e2 = op(e2, e2)) | ~ spl15_31), inference(backward_demodulation, [], [f170, f441])).
fof(f441, plain, ((e2 = op(e2, e0)) | ~ spl15_31), inference(avatar_component_clause, [], [f439])).
fof(f3114, plain, (~ spl15_3 | ~ spl15_35), inference(avatar_split_clause, [], [f3111, f456, f320])).
fof(f3111, plain, (~ (e2 = op(e3, e3)) | ~ spl15_35), inference(backward_demodulation, [], [f155, f458])).
fof(f458, plain, ((e2 = op(e1, e3)) | ~ spl15_35), inference(avatar_component_clause, [], [f456])).
fof(f155, plain, ~ (op(e1, e3) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f3109, plain, (~ spl15_33 | ~ spl15_37), inference(avatar_split_clause, [], [f3104, f465, f448])).
fof(f448, plain, (spl15_33 <=> (e0 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl15_33])])).
fof(f3104, plain, (~ (e0 = op(e1, e3)) | ~ spl15_37), inference(backward_demodulation, [], [f168, f467])).
fof(f168, plain, ~ (op(e1, e2) = op(e1, e3)), inference(cnf_transformation, [], [f3])).
fof(f3101, plain, (~ spl15_33 | ~ spl15_44 | spl15_115), inference(avatar_split_clause, [], [f3097, f932, f494, f448])).
fof(f3097, plain, (~ (e0 = op(e1, e3)) | (~ spl15_44 | spl15_115)), inference(backward_demodulation, [], [f934, f496])).
fof(f3099, plain, (~ spl15_40 | ~ spl15_44), inference(avatar_split_clause, [], [f3093, f494, f477])).
fof(f3093, plain, (~ (e3 = op(e1, e2)) | ~ spl15_44), inference(backward_demodulation, [], [f166, f496])).
fof(f3090, plain, (~ spl15_33 | ~ spl15_49), inference(avatar_split_clause, [], [f3083, f516, f448])).
fof(f3083, plain, (~ (e0 = op(e1, e3)) | ~ spl15_49), inference(backward_demodulation, [], [f151, f518])).
fof(f3082, plain, (~ spl15_53 | ~ spl15_55 | spl15_84), inference(avatar_split_clause, [], [f3080, f789, f541, f533])).
fof(f3080, plain, (~ (e0 = op(e0, e2)) | (~ spl15_55 | spl15_84)), inference(backward_demodulation, [], [f790, f543])).
fof(f790, plain, (~ (e0 = op(e0, op(e0, e2))) | spl15_84), inference(avatar_component_clause, [], [f789])).
fof(f3074, plain, (~ spl15_61 | ~ spl15_64), inference(avatar_contradiction_clause, [], [f3073])).
fof(f3073, plain, ($false | (~ spl15_61 | ~ spl15_64)), inference(subsumption_resolution, [], [f3072, f183])).
fof(f183, plain, ~ (e0 = e3), inference(cnf_transformation, [], [f4])).
fof(f3072, plain, ((e0 = e3) | (~ spl15_61 | ~ spl15_64)), inference(backward_demodulation, [], [f581, f569])).
fof(f3069, plain, (spl15_49 | ~ spl15_64 | ~ spl15_92), inference(avatar_split_clause, [], [f3066, f823, f579, f516])).
fof(f3066, plain, ((e0 = op(e0, e3)) | (~ spl15_64 | ~ spl15_92)), inference(backward_demodulation, [], [f825, f581])).
fof(f3067, plain, (~ spl15_16 | ~ spl15_64), inference(avatar_split_clause, [], [f3063, f579, f375])).
fof(f3063, plain, (~ (e3 = op(e3, e0)) | ~ spl15_64), inference(backward_demodulation, [], [f135, f581])).
fof(f3035, plain, (~ spl15_35 | ~ spl15_100 | spl15_112), inference(avatar_split_clause, [], [f3034, f916, f858, f456])).
fof(f858, plain, (spl15_100 <=> (e1 = op(e3, op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl15_100])])).
fof(f916, plain, (spl15_112 <=> (e2 = op(op(e3, op(e3, e3)), e3))), introduced(avatar_definition, [new_symbols(naming, [spl15_112])])).
fof(f3034, plain, (~ (e2 = op(e1, e3)) | (~ spl15_100 | spl15_112)), inference(forward_demodulation, [], [f918, f859])).
fof(f859, plain, ((e1 = op(e3, op(e3, e3))) | ~ spl15_100), inference(avatar_component_clause, [], [f858])).
fof(f918, plain, (~ (e2 = op(op(e3, op(e3, e3)), e3)) | spl15_112), inference(avatar_component_clause, [], [f916])).
fof(f3029, plain, (~ spl15_3 | ~ spl15_6 | ~ spl15_33 | spl15_96), inference(avatar_contradiction_clause, [], [f3028])).
fof(f3028, plain, ($false | (~ spl15_3 | ~ spl15_6 | ~ spl15_33 | spl15_96)), inference(subsumption_resolution, [], [f3027, f450])).
fof(f450, plain, ((e0 = op(e1, e3)) | ~ spl15_33), inference(avatar_component_clause, [], [f448])).
fof(f3027, plain, (~ (e0 = op(e1, e3)) | (~ spl15_3 | ~ spl15_6 | spl15_96)), inference(forward_demodulation, [], [f3026, f335])).
fof(f3026, plain, (~ (e0 = op(op(e3, e2), e3)) | (~ spl15_3 | spl15_96)), inference(forward_demodulation, [], [f842, f322])).
fof(f3022, plain, (~ spl15_40 | ~ spl15_43 | spl15_102), inference(avatar_contradiction_clause, [], [f3021])).
fof(f3021, plain, ($false | (~ spl15_40 | ~ spl15_43 | spl15_102)), inference(subsumption_resolution, [], [f3020, f479])).
fof(f479, plain, ((e3 = op(e1, e2)) | ~ spl15_40), inference(avatar_component_clause, [], [f477])).
fof(f3020, plain, (~ (e3 = op(e1, e2)) | (~ spl15_43 | spl15_102)), inference(forward_demodulation, [], [f869, f492])).
fof(f869, plain, (~ (e3 = op(e1, op(e1, e1))) | spl15_102), inference(avatar_component_clause, [], [f867])).
fof(f867, plain, (spl15_102 <=> (e3 = op(e1, op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl15_102])])).
fof(f3013, plain, (~ spl15_3 | ~ spl15_6 | spl15_100), inference(avatar_contradiction_clause, [], [f3012])).
fof(f3012, plain, ($false | (~ spl15_3 | ~ spl15_6 | spl15_100)), inference(subsumption_resolution, [], [f3009, f335])).
fof(f3009, plain, (~ (e1 = op(e3, e2)) | (~ spl15_3 | spl15_100)), inference(backward_demodulation, [], [f860, f322])).
fof(f860, plain, (~ (e1 = op(e3, op(e3, e3))) | spl15_100), inference(avatar_component_clause, [], [f858])).
fof(f3001, plain, (~ spl15_1 | ~ spl15_9), inference(avatar_split_clause, [], [f2999, f346, f312])).
fof(f2999, plain, (~ (e0 = op(e3, e3)) | ~ spl15_9), inference(backward_demodulation, [], [f179, f348])).
fof(f348, plain, ((e0 = op(e3, e1)) | ~ spl15_9), inference(avatar_component_clause, [], [f346])).
fof(f2994, plain, (~ spl15_12 | ~ spl15_16), inference(avatar_split_clause, [], [f2990, f375, f358])).
fof(f2990, plain, (~ (e3 = op(e3, e1)) | ~ spl15_16), inference(backward_demodulation, [], [f175, f377])).
fof(f2986, plain, (~ spl15_1 | ~ spl15_33), inference(avatar_split_clause, [], [f2981, f448, f312])).
fof(f2981, plain, (~ (e0 = op(e3, e3)) | ~ spl15_33), inference(backward_demodulation, [], [f155, f450])).
fof(f2979, plain, (~ spl15_9 | ~ spl15_40 | ~ spl15_43 | spl15_101), inference(avatar_split_clause, [], [f2977, f863, f490, f477, f346])).
fof(f2977, plain, (~ (e0 = op(e3, e1)) | (~ spl15_40 | ~ spl15_43 | spl15_101)), inference(backward_demodulation, [], [f2831, f479])).
fof(f2831, plain, (~ (e0 = op(op(e1, e2), e1)) | (~ spl15_43 | spl15_101)), inference(backward_demodulation, [], [f865, f492])).
fof(f2978, plain, (~ spl15_36 | ~ spl15_40), inference(avatar_split_clause, [], [f2975, f477, f460])).
fof(f2975, plain, (~ (e3 = op(e1, e3)) | ~ spl15_40), inference(backward_demodulation, [], [f168, f479])).
fof(f2972, plain, (~ spl15_14 | ~ spl15_46), inference(avatar_split_clause, [], [f2966, f503, f367])).
fof(f2966, plain, (~ (e1 = op(e3, e0)) | ~ spl15_46), inference(backward_demodulation, [], [f137, f505])).
fof(f2960, plain, (~ spl15_37 | ~ spl15_53), inference(avatar_split_clause, [], [f2951, f533, f465])).
fof(f2951, plain, (~ (e0 = op(e1, e2)) | ~ spl15_53), inference(backward_demodulation, [], [f145, f535])).
fof(f2950, plain, (~ spl15_10 | ~ spl15_58), inference(avatar_split_clause, [], [f2947, f554, f350])).
fof(f2947, plain, (~ (e1 = op(e3, e1)) | ~ spl15_58), inference(backward_demodulation, [], [f141, f556])).
fof(f2914, plain, (~ spl15_1 | ~ spl15_14 | spl15_100), inference(avatar_contradiction_clause, [], [f2913])).
fof(f2913, plain, ($false | (~ spl15_1 | ~ spl15_14 | spl15_100)), inference(subsumption_resolution, [], [f2907, f369])).
fof(f2907, plain, (~ (e1 = op(e3, e0)) | (~ spl15_1 | spl15_100)), inference(backward_demodulation, [], [f860, f314])).
fof(f314, plain, ((e0 = op(e3, e3)) | ~ spl15_1), inference(avatar_component_clause, [], [f312])).
fof(f2911, plain, (~ spl15_1 | ~ spl15_14 | spl15_33 | ~ spl15_96), inference(avatar_contradiction_clause, [], [f2910])).
fof(f2910, plain, ($false | (~ spl15_1 | ~ spl15_14 | spl15_33 | ~ spl15_96)), inference(subsumption_resolution, [], [f2909, f449])).
fof(f449, plain, (~ (e0 = op(e1, e3)) | spl15_33), inference(avatar_component_clause, [], [f448])).
fof(f2909, plain, ((e0 = op(e1, e3)) | (~ spl15_1 | ~ spl15_14 | ~ spl15_96)), inference(forward_demodulation, [], [f2905, f369])).
fof(f2905, plain, ((e0 = op(op(e3, e0), e3)) | (~ spl15_1 | ~ spl15_96)), inference(backward_demodulation, [], [f841, f314])).
fof(f841, plain, ((e0 = op(op(e3, op(e3, e3)), e3)) | ~ spl15_96), inference(avatar_component_clause, [], [f840])).
fof(f2900, plain, (~ spl15_3 | ~ spl15_7), inference(avatar_split_clause, [], [f2897, f337, f320])).
fof(f2897, plain, (~ (e2 = op(e3, e3)) | ~ spl15_7), inference(backward_demodulation, [], [f180, f339])).
fof(f339, plain, ((e2 = op(e3, e2)) | ~ spl15_7), inference(avatar_component_clause, [], [f337])).
fof(f2896, plain, (~ spl15_10 | ~ spl15_12), inference(avatar_contradiction_clause, [], [f2895])).
fof(f2895, plain, ($false | (~ spl15_10 | ~ spl15_12)), inference(subsumption_resolution, [], [f2894, f185])).
fof(f2894, plain, ((e1 = e3) | (~ spl15_10 | ~ spl15_12)), inference(forward_demodulation, [], [f360, f352])).
fof(f2878, plain, (~ spl15_22 | ~ spl15_23), inference(avatar_contradiction_clause, [], [f2877])).
fof(f2877, plain, ($false | (~ spl15_22 | ~ spl15_23)), inference(subsumption_resolution, [], [f2876, f184])).
fof(f2876, plain, ((e1 = e2) | (~ spl15_22 | ~ spl15_23)), inference(backward_demodulation, [], [f407, f403])).
fof(f2871, plain, (~ spl15_23 | spl15_86), inference(avatar_contradiction_clause, [], [f2870])).
fof(f2870, plain, ($false | (~ spl15_23 | spl15_86)), inference(subsumption_resolution, [], [f2864, f407])).
fof(f2864, plain, (~ (e2 = op(e2, e2)) | (~ spl15_23 | spl15_86)), inference(backward_demodulation, [], [f798, f407])).
fof(f798, plain, (~ (e2 = op(e2, op(e2, e2))) | spl15_86), inference(avatar_component_clause, [], [f797])).
fof(f797, plain, (spl15_86 <=> (e2 = op(e2, op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl15_86])])).
fof(f2859, plain, (~ spl15_24 | ~ spl15_28), inference(avatar_split_clause, [], [f2856, f426, f409])).
fof(f2856, plain, (~ (e3 = op(e2, e2)) | ~ spl15_28), inference(backward_demodulation, [], [f172, f428])).
fof(f2841, plain, (~ spl15_42 | ~ spl15_43), inference(avatar_contradiction_clause, [], [f2840])).
fof(f2840, plain, ($false | (~ spl15_42 | ~ spl15_43)), inference(subsumption_resolution, [], [f2839, f184])).
fof(f2839, plain, ((e1 = e2) | (~ spl15_42 | ~ spl15_43)), inference(backward_demodulation, [], [f492, f488])).
fof(f488, plain, ((e1 = op(e1, e1)) | ~ spl15_42), inference(avatar_component_clause, [], [f486])).
fof(f2835, plain, (~ spl15_27 | ~ spl15_43), inference(avatar_split_clause, [], [f2827, f490, f422])).
fof(f2827, plain, (~ (e2 = op(e2, e1)) | ~ spl15_43), inference(backward_demodulation, [], [f142, f492])).
fof(f2826, plain, (~ spl15_16 | ~ spl15_48), inference(avatar_split_clause, [], [f2825, f511, f375])).
fof(f2825, plain, (~ (e3 = op(e3, e0)) | ~ spl15_48), inference(backward_demodulation, [], [f137, f513])).
fof(f2821, plain, (~ spl15_18 | ~ spl15_50), inference(avatar_split_clause, [], [f2816, f520, f384])).
fof(f2816, plain, (~ (e1 = op(e2, e3)) | ~ spl15_50), inference(backward_demodulation, [], [f152, f522])).
fof(f2805, plain, (~ spl15_41 | ~ spl15_57), inference(avatar_split_clause, [], [f2797, f550, f482])).
fof(f2797, plain, (~ (e0 = op(e1, e1)) | ~ spl15_57), inference(backward_demodulation, [], [f139, f552])).
fof(f2790, plain, (~ spl15_56 | ~ spl15_63 | spl15_109), inference(avatar_split_clause, [], [f2789, f901, f575, f545])).
fof(f2789, plain, (~ (e3 = op(e0, e2)) | (~ spl15_63 | spl15_109)), inference(forward_demodulation, [], [f903, f577])).
fof(f903, plain, (~ (e3 = op(e0, op(e0, e0))) | spl15_109), inference(avatar_component_clause, [], [f901])).
fof(f2779, plain, (~ spl15_42 | ~ spl15_10), inference(avatar_split_clause, [], [f2778, f350, f486])).
fof(f2778, plain, (~ (e1 = op(e1, e1)) | ~ spl15_10), inference(forward_demodulation, [], [f143, f352])).
fof(f2771, plain, (~ spl15_14 | ~ spl15_10), inference(avatar_split_clause, [], [f2770, f350, f367])).
fof(f2770, plain, (~ (e1 = op(e3, e0)) | ~ spl15_10), inference(forward_demodulation, [], [f175, f352])).
fof(f2750, plain, (~ spl15_6 | ~ spl15_10), inference(avatar_split_clause, [], [f2745, f350, f333])).
fof(f2745, plain, (~ (e1 = op(e3, e2)) | ~ spl15_10), inference(backward_demodulation, [], [f178, f352])).
fof(f2744, plain, (spl15_4 | ~ spl15_16 | ~ spl15_95), inference(avatar_split_clause, [], [f2742, f835, f375, f324])).
fof(f2742, plain, ((e3 = op(e3, e3)) | (~ spl15_16 | ~ spl15_95)), inference(backward_demodulation, [], [f837, f377])).
fof(f837, plain, ((e3 = op(e3, op(e3, e0))) | ~ spl15_95), inference(avatar_component_clause, [], [f835])).
fof(f2743, plain, (~ spl15_4 | ~ spl15_16), inference(avatar_split_clause, [], [f2741, f375, f324])).
fof(f2741, plain, (~ (e3 = op(e3, e3)) | ~ spl15_16), inference(backward_demodulation, [], [f177, f377])).
fof(f2738, plain, (~ spl15_4 | ~ spl15_36), inference(avatar_split_clause, [], [f2736, f460, f324])).
fof(f2736, plain, (~ (e3 = op(e3, e3)) | ~ spl15_36), inference(backward_demodulation, [], [f155, f462])).
fof(f2735, plain, (~ spl15_35 | ~ spl15_39), inference(avatar_split_clause, [], [f2732, f473, f456])).
fof(f2732, plain, (~ (e2 = op(e1, e3)) | ~ spl15_39), inference(backward_demodulation, [], [f168, f475])).
fof(f2729, plain, (~ spl15_9 | ~ spl15_41), inference(avatar_split_clause, [], [f2722, f482, f346])).
fof(f2722, plain, (~ (e0 = op(e3, e1)) | ~ spl15_41), inference(backward_demodulation, [], [f143, f484])).
fof(f2719, plain, (~ spl15_46 | ~ spl15_54 | ~ spl15_63 | spl15_108), inference(avatar_contradiction_clause, [], [f2718])).
fof(f2718, plain, ($false | (~ spl15_46 | ~ spl15_54 | ~ spl15_63 | spl15_108)), inference(subsumption_resolution, [], [f2714, f505])).
fof(f2714, plain, (~ (e1 = op(e1, e0)) | (~ spl15_54 | ~ spl15_63 | spl15_108)), inference(backward_demodulation, [], [f2697, f539])).
fof(f2697, plain, (~ (e1 = op(op(e0, e2), e0)) | (~ spl15_63 | spl15_108)), inference(backward_demodulation, [], [f899, f577])).
fof(f2717, plain, (~ spl15_6 | ~ spl15_54), inference(avatar_split_clause, [], [f2710, f537, f333])).
fof(f2710, plain, (~ (e1 = op(e3, e2)) | ~ spl15_54), inference(backward_demodulation, [], [f147, f539])).
fof(f2702, plain, (~ spl15_55 | ~ spl15_63), inference(avatar_split_clause, [], [f2694, f575, f541])).
fof(f2694, plain, (~ (e2 = op(e0, e2)) | ~ spl15_63), inference(backward_demodulation, [], [f158, f577])).
fof(f158, plain, ~ (op(e0, e0) = op(e0, e2)), inference(cnf_transformation, [], [f3])).
fof(f2684, plain, (~ spl15_37 | ~ spl15_18 | ~ spl15_24 | spl15_98), inference(avatar_split_clause, [], [f2683, f849, f409, f384, f465])).
fof(f2683, plain, (~ (e0 = op(e1, e2)) | (~ spl15_18 | ~ spl15_24 | spl15_98)), inference(forward_demodulation, [], [f2682, f386])).
fof(f2682, plain, (~ (e0 = op(op(e2, e3), e2)) | (~ spl15_24 | spl15_98)), inference(forward_demodulation, [], [f851, f411])).
fof(f2670, plain, (spl15_1 | ~ spl15_4 | ~ spl15_107), inference(avatar_contradiction_clause, [], [f2669])).
fof(f2669, plain, ($false | (spl15_1 | ~ spl15_4 | ~ spl15_107)), inference(subsumption_resolution, [], [f2668, f313])).
fof(f313, plain, (~ (e0 = op(e3, e3)) | spl15_1), inference(avatar_component_clause, [], [f312])).
fof(f2668, plain, ((e0 = op(e3, e3)) | (~ spl15_4 | ~ spl15_107)), inference(backward_demodulation, [], [f893, f326])).
fof(f326, plain, ((e3 = op(e3, e3)) | ~ spl15_4), inference(avatar_component_clause, [], [f324])).
fof(f893, plain, ((e0 = op(e3, op(e3, e3))) | ~ spl15_107), inference(avatar_component_clause, [], [f892])).
fof(f2631, plain, (~ spl15_39 | ~ spl15_55), inference(avatar_split_clause, [], [f2626, f541, f473])).
fof(f2626, plain, (~ (e2 = op(e1, e2)) | ~ spl15_55), inference(backward_demodulation, [], [f145, f543])).
fof(f2615, plain, (spl15_23 | ~ spl15_27 | ~ spl15_90), inference(avatar_contradiction_clause, [], [f2614])).
fof(f2614, plain, ($false | (spl15_23 | ~ spl15_27 | ~ spl15_90)), inference(subsumption_resolution, [], [f2613, f406])).
fof(f406, plain, (~ (e2 = op(e2, e2)) | spl15_23), inference(avatar_component_clause, [], [f405])).
fof(f2613, plain, ((e2 = op(e2, e2)) | (~ spl15_27 | ~ spl15_90)), inference(forward_demodulation, [], [f816, f424])).
fof(f2610, plain, (~ spl15_99 | ~ spl15_103), inference(avatar_contradiction_clause, [], [f2609])).
fof(f2609, plain, ($false | (~ spl15_99 | ~ spl15_103)), inference(subsumption_resolution, [], [f2608, f185])).
fof(f2608, plain, ((e1 = e3) | (~ spl15_99 | ~ spl15_103)), inference(backward_demodulation, [], [f854, f873])).
fof(f873, plain, ((e1 = op(e2, op(e2, e2))) | ~ spl15_103), inference(avatar_component_clause, [], [f872])).
fof(f872, plain, (spl15_103 <=> (e1 = op(e2, op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl15_103])])).
fof(f2573, plain, (~ spl15_6 | ~ spl15_9 | ~ spl15_87), inference(avatar_contradiction_clause, [], [f2572])).
fof(f2572, plain, ($false | (~ spl15_6 | ~ spl15_9 | ~ spl15_87)), inference(subsumption_resolution, [], [f2571, f183])).
fof(f2571, plain, ((e0 = e3) | (~ spl15_6 | ~ spl15_9 | ~ spl15_87)), inference(forward_demodulation, [], [f2569, f348])).
fof(f2569, plain, ((e3 = op(e3, e1)) | (~ spl15_6 | ~ spl15_87)), inference(backward_demodulation, [], [f803, f335])).
fof(f803, plain, ((e3 = op(e3, op(e3, e2))) | ~ spl15_87), inference(avatar_component_clause, [], [f801])).
fof(f801, plain, (spl15_87 <=> (e3 = op(e3, op(e3, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl15_87])])).
fof(f2557, plain, (~ spl15_2 | ~ spl15_18), inference(avatar_split_clause, [], [f2553, f384, f316])).
fof(f2553, plain, (~ (e1 = op(e3, e3)) | ~ spl15_18), inference(backward_demodulation, [], [f156, f386])).
fof(f2549, plain, (~ spl15_22 | ~ spl15_27 | spl15_86), inference(avatar_contradiction_clause, [], [f2548])).
fof(f2548, plain, ($false | (~ spl15_22 | ~ spl15_27 | spl15_86)), inference(subsumption_resolution, [], [f2542, f424])).
fof(f2542, plain, (~ (e2 = op(e2, e1)) | (~ spl15_22 | spl15_86)), inference(backward_demodulation, [], [f798, f403])).
fof(f2537, plain, (~ spl15_27 | ~ spl15_28), inference(avatar_contradiction_clause, [], [f2536])).
fof(f2536, plain, ($false | (~ spl15_27 | ~ spl15_28)), inference(subsumption_resolution, [], [f2535, f186])).
fof(f186, plain, ~ (e2 = e3), inference(cnf_transformation, [], [f4])).
fof(f2535, plain, ((e2 = e3) | (~ spl15_27 | ~ spl15_28)), inference(forward_demodulation, [], [f428, f424])).
fof(f2533, plain, (~ spl15_21 | ~ spl15_29), inference(avatar_split_clause, [], [f2528, f431, f397])).
fof(f2528, plain, (~ (e0 = op(e2, e2)) | ~ spl15_29), inference(backward_demodulation, [], [f170, f433])).
fof(f2519, plain, (~ spl15_39 | ~ spl15_40), inference(avatar_contradiction_clause, [], [f2518])).
fof(f2518, plain, ($false | (~ spl15_39 | ~ spl15_40)), inference(subsumption_resolution, [], [f2517, f186])).
fof(f2517, plain, ((e2 = e3) | (~ spl15_39 | ~ spl15_40)), inference(backward_demodulation, [], [f479, f475])).
fof(f2512, plain, (spl15_38 | ~ spl15_43 | ~ spl15_89), inference(avatar_split_clause, [], [f2508, f810, f490, f469])).
fof(f2508, plain, ((e1 = op(e1, e2)) | (~ spl15_43 | ~ spl15_89)), inference(backward_demodulation, [], [f812, f492])).
fof(f2494, plain, (~ spl15_20 | ~ spl15_52), inference(avatar_split_clause, [], [f2491, f528, f392])).
fof(f2491, plain, (~ (e3 = op(e2, e3)) | ~ spl15_52), inference(backward_demodulation, [], [f152, f530])).
fof(f2489, plain, (~ spl15_21 | ~ spl15_53), inference(avatar_split_clause, [], [f2480, f533, f397])).
fof(f2480, plain, (~ (e0 = op(e2, e2)) | ~ spl15_53), inference(backward_demodulation, [], [f146, f535])).
fof(f2459, plain, (~ spl15_50 | spl15_105 | ~ spl15_107), inference(avatar_split_clause, [], [f2458, f892, f882, f520])).
fof(f2458, plain, (~ (e1 = op(e0, e3)) | (spl15_105 | ~ spl15_107)), inference(forward_demodulation, [], [f884, f893])).
fof(f2439, plain, (~ spl15_31 | ~ spl15_63), inference(avatar_split_clause, [], [f2438, f575, f439])).
fof(f2438, plain, (~ (e2 = op(e2, e0)) | ~ spl15_63), inference(forward_demodulation, [], [f134, f577])).
fof(f134, plain, ~ (op(e0, e0) = op(e2, e0)), inference(cnf_transformation, [], [f3])).
fof(f2437, plain, (~ spl15_23 | ~ spl15_27), inference(avatar_split_clause, [], [f2436, f422, f405])).
fof(f2436, plain, (~ (e2 = op(e2, e2)) | ~ spl15_27), inference(forward_demodulation, [], [f172, f424])).
fof(f2432, plain, (~ spl15_11 | ~ spl15_27), inference(avatar_split_clause, [], [f2431, f422, f354])).
fof(f2431, plain, (~ (e2 = op(e3, e1)) | ~ spl15_27), inference(forward_demodulation, [], [f144, f424])).
fof(f2428, plain, (~ spl15_7 | ~ spl15_87), inference(avatar_contradiction_clause, [], [f2427])).
fof(f2427, plain, ($false | (~ spl15_7 | ~ spl15_87)), inference(subsumption_resolution, [], [f2426, f186])).
fof(f2426, plain, ((e2 = e3) | (~ spl15_7 | ~ spl15_87)), inference(forward_demodulation, [], [f2425, f339])).
fof(f2425, plain, ((e3 = op(e3, e2)) | (~ spl15_7 | ~ spl15_87)), inference(backward_demodulation, [], [f803, f339])).
fof(f2423, plain, (~ spl15_1 | ~ spl15_16 | spl15_96), inference(avatar_contradiction_clause, [], [f2422])).
fof(f2422, plain, ($false | (~ spl15_1 | ~ spl15_16 | spl15_96)), inference(subsumption_resolution, [], [f2414, f314])).
fof(f2414, plain, (~ (e0 = op(e3, e3)) | (~ spl15_1 | ~ spl15_16 | spl15_96)), inference(backward_demodulation, [], [f2359, f377])).
fof(f2359, plain, (~ (e0 = op(op(e3, e0), e3)) | (~ spl15_1 | spl15_96)), inference(forward_demodulation, [], [f842, f314])).
fof(f2401, plain, (~ spl15_19 | ~ spl15_35), inference(avatar_split_clause, [], [f2398, f456, f388])).
fof(f2398, plain, (~ (e2 = op(e2, e3)) | ~ spl15_35), inference(backward_demodulation, [], [f154, f458])).
fof(f2396, plain, (~ spl15_36 | ~ spl15_44), inference(avatar_split_clause, [], [f2391, f494, f460])).
fof(f2391, plain, (~ (e3 = op(e1, e3)) | ~ spl15_44), inference(backward_demodulation, [], [f167, f496])).
fof(f2367, plain, (~ spl15_10 | ~ spl15_91), inference(avatar_contradiction_clause, [], [f2366])).
fof(f2366, plain, ($false | (~ spl15_10 | ~ spl15_91)), inference(subsumption_resolution, [], [f2365, f185])).
fof(f2365, plain, ((e1 = e3) | (~ spl15_10 | ~ spl15_91)), inference(forward_demodulation, [], [f2364, f352])).
fof(f2364, plain, ((e3 = op(e3, e1)) | (~ spl15_10 | ~ spl15_91)), inference(forward_demodulation, [], [f820, f352])).
fof(f2344, plain, (~ spl15_27 | ~ spl15_30 | spl15_94), inference(avatar_split_clause, [], [f2293, f831, f435, f422])).
fof(f2293, plain, (~ (e2 = op(e2, e1)) | (~ spl15_30 | spl15_94)), inference(backward_demodulation, [], [f832, f437])).
fof(f2337, plain, (~ spl15_15 | ~ spl15_1 | spl15_97), inference(avatar_split_clause, [], [f2328, f844, f312, f371])).
fof(f2328, plain, (~ (e2 = op(e3, e0)) | (~ spl15_1 | spl15_97)), inference(backward_demodulation, [], [f846, f314])).
fof(f846, plain, (~ (e2 = op(e3, op(e3, e3))) | spl15_97), inference(avatar_component_clause, [], [f844])).
fof(f2333, plain, (~ spl15_1 | ~ spl15_15 | ~ spl15_83), inference(avatar_contradiction_clause, [], [f2332])).
fof(f2332, plain, ($false | (~ spl15_1 | ~ spl15_15 | ~ spl15_83)), inference(subsumption_resolution, [], [f2331, f186])).
fof(f2331, plain, ((e2 = e3) | (~ spl15_1 | ~ spl15_15 | ~ spl15_83)), inference(forward_demodulation, [], [f2327, f373])).
fof(f2327, plain, ((e3 = op(e3, e0)) | (~ spl15_1 | ~ spl15_83)), inference(backward_demodulation, [], [f786, f314])).
fof(f2323, plain, (spl15_4 | ~ spl15_8 | ~ spl15_87), inference(avatar_split_clause, [], [f2321, f801, f341, f324])).
fof(f2321, plain, ((e3 = op(e3, e3)) | (~ spl15_8 | ~ spl15_87)), inference(backward_demodulation, [], [f803, f343])).
fof(f2313, plain, (spl15_8 | ~ spl15_15 | ~ spl15_95), inference(avatar_split_clause, [], [f2311, f835, f371, f341])).
fof(f2311, plain, ((e3 = op(e3, e2)) | (~ spl15_15 | ~ spl15_95)), inference(backward_demodulation, [], [f837, f373])).
fof(f2312, plain, (~ spl15_11 | ~ spl15_15), inference(avatar_split_clause, [], [f2307, f371, f354])).
fof(f2307, plain, (~ (e2 = op(e3, e1)) | ~ spl15_15), inference(backward_demodulation, [], [f175, f373])).
fof(f2306, plain, (~ spl15_5 | ~ spl15_21), inference(avatar_split_clause, [], [f2300, f397, f329])).
fof(f2300, plain, (~ (e0 = op(e3, e2)) | ~ spl15_21), inference(backward_demodulation, [], [f150, f399])).
fof(f2294, plain, (~ spl15_14 | ~ spl15_30), inference(avatar_split_clause, [], [f2288, f435, f367])).
fof(f2288, plain, (~ (e1 = op(e3, e0)) | ~ spl15_30), inference(backward_demodulation, [], [f138, f437])).
fof(f2267, plain, (~ spl15_33 | ~ spl15_45), inference(avatar_split_clause, [], [f2261, f499, f448])).
fof(f2261, plain, (~ (e0 = op(e1, e3)) | ~ spl15_45), inference(backward_demodulation, [], [f165, f501])).
fof(f2246, plain, (~ spl15_1 | ~ spl15_83 | spl15_96), inference(avatar_split_clause, [], [f2244, f840, f784, f312])).
fof(f2244, plain, (~ (e0 = op(e3, e3)) | (~ spl15_83 | spl15_96)), inference(backward_demodulation, [], [f842, f786])).
fof(f2243, plain, (~ spl15_49 | ~ spl15_64 | spl15_92), inference(avatar_split_clause, [], [f2242, f823, f579, f516])).
fof(f2242, plain, (~ (e0 = op(e0, e3)) | (~ spl15_64 | spl15_92)), inference(forward_demodulation, [], [f824, f581])).
fof(f2233, plain, (~ spl15_50 | ~ spl15_64 | spl15_116), inference(avatar_contradiction_clause, [], [f2232])).
fof(f2232, plain, ($false | (~ spl15_50 | ~ spl15_64 | spl15_116)), inference(subsumption_resolution, [], [f2231, f522])).
fof(f2231, plain, (~ (e1 = op(e0, e3)) | (~ spl15_64 | spl15_116)), inference(forward_demodulation, [], [f939, f581])).
fof(f939, plain, (~ (e1 = op(e0, op(e0, e0))) | spl15_116), inference(avatar_component_clause, [], [f937])).
fof(f937, plain, (spl15_116 <=> (e1 = op(e0, op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl15_116])])).
fof(f2228, plain, (~ spl15_47 | ~ spl15_50 | ~ spl15_64 | spl15_114), inference(avatar_split_clause, [], [f2227, f927, f579, f520, f507])).
fof(f927, plain, (spl15_114 <=> (e2 = op(op(e0, op(e0, e0)), e0))), introduced(avatar_definition, [new_symbols(naming, [spl15_114])])).
fof(f2227, plain, (~ (e2 = op(e1, e0)) | (~ spl15_50 | ~ spl15_64 | spl15_114)), inference(forward_demodulation, [], [f2226, f522])).
fof(f2226, plain, (~ (e2 = op(op(e0, e3), e0)) | (~ spl15_64 | spl15_114)), inference(forward_demodulation, [], [f929, f581])).
fof(f929, plain, (~ (e2 = op(op(e0, op(e0, e0)), e0)) | spl15_114), inference(avatar_component_clause, [], [f927])).
fof(f2216, plain, (~ spl15_60 | ~ spl15_115 | spl15_118), inference(avatar_split_clause, [], [f2215, f947, f932, f562])).
fof(f2215, plain, (~ (e3 = op(e0, e1)) | (~ spl15_115 | spl15_118)), inference(forward_demodulation, [], [f949, f933])).
fof(f2208, plain, (~ spl15_32 | ~ spl15_64), inference(avatar_split_clause, [], [f2207, f579, f443])).
fof(f2207, plain, (~ (e3 = op(e2, e0)) | ~ spl15_64), inference(forward_demodulation, [], [f134, f581])).
fof(f2204, plain, (~ spl15_7 | ~ spl15_55), inference(avatar_split_clause, [], [f2203, f541, f337])).
fof(f2203, plain, (~ (e2 = op(e3, e2)) | ~ spl15_55), inference(forward_demodulation, [], [f147, f543])).
fof(f2196, plain, (~ spl15_2 | ~ spl15_14), inference(avatar_split_clause, [], [f2193, f367, f316])).
fof(f2193, plain, (~ (e1 = op(e3, e3)) | ~ spl15_14), inference(backward_demodulation, [], [f177, f369])).
fof(f2190, plain, (~ spl15_22 | ~ spl15_24), inference(avatar_contradiction_clause, [], [f2189])).
fof(f2189, plain, ($false | (~ spl15_22 | ~ spl15_24)), inference(subsumption_resolution, [], [f2188, f185])).
fof(f2188, plain, ((e1 = e3) | (~ spl15_22 | ~ spl15_24)), inference(forward_demodulation, [], [f411, f403])).
fof(f2184, plain, (~ spl15_22 | ~ spl15_25 | spl15_54 | ~ spl15_106), inference(avatar_contradiction_clause, [], [f2183])).
fof(f2183, plain, ($false | (~ spl15_22 | ~ spl15_25 | spl15_54 | ~ spl15_106)), inference(subsumption_resolution, [], [f2179, f538])).
fof(f538, plain, (~ (e1 = op(e0, e2)) | spl15_54), inference(avatar_component_clause, [], [f537])).
fof(f2179, plain, ((e1 = op(e0, e2)) | (~ spl15_22 | ~ spl15_25 | ~ spl15_106)), inference(backward_demodulation, [], [f2142, f416])).
fof(f2142, plain, ((e1 = op(op(e2, e1), e2)) | (~ spl15_22 | ~ spl15_106)), inference(forward_demodulation, [], [f888, f403])).
fof(f888, plain, ((e1 = op(op(e2, op(e2, e2)), e2)) | ~ spl15_106), inference(avatar_component_clause, [], [f887])).
fof(f2175, plain, (~ spl15_46 | ~ spl15_47), inference(avatar_contradiction_clause, [], [f2174])).
fof(f2174, plain, ($false | (~ spl15_46 | ~ spl15_47)), inference(subsumption_resolution, [], [f2173, f184])).
fof(f2173, plain, ((e1 = e2) | (~ spl15_46 | ~ spl15_47)), inference(forward_demodulation, [], [f509, f505])).
fof(f2172, plain, (~ spl15_49 | ~ spl15_50), inference(avatar_contradiction_clause, [], [f2171])).
fof(f2171, plain, ($false | (~ spl15_49 | ~ spl15_50)), inference(subsumption_resolution, [], [f2170, f181])).
fof(f181, plain, ~ (e0 = e1), inference(cnf_transformation, [], [f4])).
fof(f2170, plain, ((e0 = e1) | (~ spl15_49 | ~ spl15_50)), inference(backward_demodulation, [], [f522, f518])).
fof(f2157, plain, (~ spl15_56 | ~ spl15_64), inference(avatar_split_clause, [], [f2153, f579, f545])).
fof(f2153, plain, (~ (e3 = op(e0, e2)) | ~ spl15_64), inference(backward_demodulation, [], [f158, f581])).
fof(f2148, plain, (~ spl15_28 | ~ spl15_22 | spl15_99), inference(avatar_split_clause, [], [f2147, f853, f401, f426])).
fof(f2147, plain, (~ (e3 = op(e2, e1)) | (~ spl15_22 | spl15_99)), inference(forward_demodulation, [], [f855, f403])).
fof(f2120, plain, (~ spl15_25 | ~ spl15_22 | spl15_110), inference(avatar_split_clause, [], [f2104, f906, f401, f414])).
fof(f2104, plain, (~ (e0 = op(e2, e1)) | (~ spl15_22 | spl15_110)), inference(backward_demodulation, [], [f908, f403])).
fof(f2118, plain, (spl15_16 | ~ spl15_5 | ~ spl15_87), inference(avatar_split_clause, [], [f1940, f801, f329, f375])).
fof(f1940, plain, ((e3 = op(e3, e0)) | (~ spl15_5 | ~ spl15_87)), inference(forward_demodulation, [], [f803, f331])).
fof(f2106, plain, (~ spl15_22 | spl15_27 | ~ spl15_86), inference(avatar_contradiction_clause, [], [f2105])).
fof(f2105, plain, ($false | (~ spl15_22 | spl15_27 | ~ spl15_86)), inference(subsumption_resolution, [], [f2102, f423])).
fof(f423, plain, (~ (e2 = op(e2, e1)) | spl15_27), inference(avatar_component_clause, [], [f422])).
fof(f2102, plain, ((e2 = op(e2, e1)) | (~ spl15_22 | ~ spl15_86)), inference(backward_demodulation, [], [f799, f403])).
fof(f799, plain, ((e2 = op(e2, op(e2, e2))) | ~ spl15_86), inference(avatar_component_clause, [], [f797])).
fof(f2085, plain, (~ spl15_57 | ~ spl15_58), inference(avatar_contradiction_clause, [], [f2084])).
fof(f2084, plain, ($false | (~ spl15_57 | ~ spl15_58)), inference(subsumption_resolution, [], [f2083, f181])).
fof(f2083, plain, ((e0 = e1) | (~ spl15_57 | ~ spl15_58)), inference(forward_demodulation, [], [f556, f552])).
fof(f2075, plain, (spl15_46 | ~ spl15_33 | ~ spl15_81), inference(avatar_split_clause, [], [f2074, f776, f448, f503])).
fof(f2074, plain, ((e1 = op(e1, e0)) | (~ spl15_33 | ~ spl15_81)), inference(forward_demodulation, [], [f778, f450])).
fof(f2044, plain, (~ spl15_52 | ~ spl15_56), inference(avatar_split_clause, [], [f2043, f545, f528])).
fof(f2043, plain, (~ (e3 = op(e0, e3)) | ~ spl15_56), inference(forward_demodulation, [], [f162, f547])).
fof(f162, plain, ~ (op(e0, e2) = op(e0, e3)), inference(cnf_transformation, [], [f3])).
fof(f2036, plain, (~ spl15_48 | ~ spl15_44), inference(avatar_split_clause, [], [f2035, f494, f511])).
fof(f2035, plain, (~ (e3 = op(e1, e0)) | ~ spl15_44), inference(forward_demodulation, [], [f163, f496])).
fof(f2003, plain, (~ spl15_55 | ~ spl15_56), inference(avatar_contradiction_clause, [], [f2002])).
fof(f2002, plain, ($false | (~ spl15_55 | ~ spl15_56)), inference(subsumption_resolution, [], [f2001, f186])).
fof(f2001, plain, ((e2 = e3) | (~ spl15_55 | ~ spl15_56)), inference(forward_demodulation, [], [f547, f543])).
fof(f1984, plain, (~ spl15_26 | ~ spl15_90), inference(avatar_contradiction_clause, [], [f1983])).
fof(f1983, plain, ($false | (~ spl15_26 | ~ spl15_90)), inference(subsumption_resolution, [], [f1982, f184])).
fof(f1982, plain, ((e1 = e2) | (~ spl15_26 | ~ spl15_90)), inference(forward_demodulation, [], [f1981, f420])).
fof(f1981, plain, ((e2 = op(e2, e1)) | (~ spl15_26 | ~ spl15_90)), inference(forward_demodulation, [], [f816, f420])).
fof(f1977, plain, (~ spl15_57 | ~ spl15_62 | spl15_92), inference(avatar_split_clause, [], [f1652, f823, f571, f550])).
fof(f1652, plain, (~ (e0 = op(e0, e1)) | (~ spl15_62 | spl15_92)), inference(backward_demodulation, [], [f824, f573])).
fof(f1975, plain, (~ spl15_51 | ~ spl15_55), inference(avatar_split_clause, [], [f1974, f541, f524])).
fof(f1974, plain, (~ (e2 = op(e0, e3)) | ~ spl15_55), inference(forward_demodulation, [], [f162, f543])).
fof(f1962, plain, (~ spl15_42 | ~ spl15_26), inference(avatar_split_clause, [], [f1783, f418, f486])).
fof(f1783, plain, (~ (e1 = op(e1, e1)) | ~ spl15_26), inference(forward_demodulation, [], [f142, f420])).
fof(f1961, plain, (~ spl15_42 | ~ spl15_38), inference(avatar_split_clause, [], [f1785, f469, f486])).
fof(f1785, plain, (~ (e1 = op(e1, e1)) | ~ spl15_38), inference(forward_demodulation, [], [f166, f471])).
fof(f1953, plain, (~ spl15_2 | ~ spl15_11 | ~ spl15_83), inference(avatar_contradiction_clause, [], [f1952])).
fof(f1952, plain, ($false | (~ spl15_2 | ~ spl15_11 | ~ spl15_83)), inference(subsumption_resolution, [], [f1951, f186])).
fof(f1951, plain, ((e2 = e3) | (~ spl15_2 | ~ spl15_11 | ~ spl15_83)), inference(forward_demodulation, [], [f1950, f356])).
fof(f1950, plain, ((e3 = op(e3, e1)) | (~ spl15_2 | ~ spl15_83)), inference(forward_demodulation, [], [f786, f318])).
fof(f318, plain, ((e1 = op(e3, e3)) | ~ spl15_2), inference(avatar_component_clause, [], [f316])).
fof(f1948, plain, (~ spl15_19 | spl15_23 | ~ spl15_82), inference(avatar_contradiction_clause, [], [f1947])).
fof(f1947, plain, ($false | (~ spl15_19 | spl15_23 | ~ spl15_82)), inference(subsumption_resolution, [], [f1946, f406])).
fof(f1946, plain, ((e2 = op(e2, e2)) | (~ spl15_19 | ~ spl15_82)), inference(forward_demodulation, [], [f782, f390])).
fof(f1944, plain, (spl15_34 | ~ spl15_36 | ~ spl15_81), inference(avatar_contradiction_clause, [], [f1943])).
fof(f1943, plain, ($false | (spl15_34 | ~ spl15_36 | ~ spl15_81)), inference(subsumption_resolution, [], [f1942, f453])).
fof(f453, plain, (~ (e1 = op(e1, e3)) | spl15_34), inference(avatar_component_clause, [], [f452])).
fof(f1942, plain, ((e1 = op(e1, e3)) | (~ spl15_36 | ~ spl15_81)), inference(backward_demodulation, [], [f778, f462])).
fof(f1927, plain, (~ spl15_35 | ~ spl15_47), inference(avatar_split_clause, [], [f1926, f507, f456])).
fof(f1926, plain, (~ (e2 = op(e1, e3)) | ~ spl15_47), inference(forward_demodulation, [], [f165, f509])).
fof(f1922, plain, (~ spl15_17 | ~ spl15_2 | ~ spl15_11 | spl15_96), inference(avatar_split_clause, [], [f1823, f840, f354, f316, f380])).
fof(f1823, plain, (~ (e0 = op(e2, e3)) | (~ spl15_2 | ~ spl15_11 | spl15_96)), inference(backward_demodulation, [], [f1795, f356])).
fof(f1795, plain, (~ (e0 = op(op(e3, e1), e3)) | (~ spl15_2 | spl15_96)), inference(forward_demodulation, [], [f842, f318])).
fof(f1921, plain, (~ spl15_17 | ~ spl15_49), inference(avatar_split_clause, [], [f1851, f516, f380])).
fof(f1851, plain, (~ (e0 = op(e2, e3)) | ~ spl15_49), inference(forward_demodulation, [], [f152, f518])).
fof(f1920, plain, (~ spl15_17 | ~ spl15_29), inference(avatar_split_clause, [], [f1888, f431, f380])).
fof(f1888, plain, (~ (e0 = op(e2, e3)) | ~ spl15_29), inference(backward_demodulation, [], [f171, f433])).
fof(f1919, plain, (~ spl15_2 | ~ spl15_4), inference(avatar_contradiction_clause, [], [f1918])).
fof(f1918, plain, ($false | (~ spl15_2 | ~ spl15_4)), inference(subsumption_resolution, [], [f1917, f185])).
fof(f1917, plain, ((e1 = e3) | (~ spl15_2 | ~ spl15_4)), inference(forward_demodulation, [], [f326, f318])).
fof(f1916, plain, (~ spl15_5 | ~ spl15_16 | spl15_87), inference(avatar_contradiction_clause, [], [f1915])).
fof(f1915, plain, ($false | (~ spl15_5 | ~ spl15_16 | spl15_87)), inference(subsumption_resolution, [], [f1912, f377])).
fof(f1912, plain, (~ (e3 = op(e3, e0)) | (~ spl15_5 | spl15_87)), inference(backward_demodulation, [], [f802, f331])).
fof(f802, plain, (~ (e3 = op(e3, op(e3, e2))) | spl15_87), inference(avatar_component_clause, [], [f801])).
fof(f1904, plain, (spl15_19 | ~ spl15_24 | ~ spl15_86), inference(avatar_split_clause, [], [f1900, f797, f409, f388])).
fof(f1900, plain, ((e2 = op(e2, e3)) | (~ spl15_24 | ~ spl15_86)), inference(backward_demodulation, [], [f799, f411])).
fof(f1903, plain, (~ spl15_20 | ~ spl15_24), inference(avatar_split_clause, [], [f1899, f409, f392])).
fof(f1899, plain, (~ (e3 = op(e2, e3)) | ~ spl15_24), inference(backward_demodulation, [], [f174, f411])).
fof(f1897, plain, (~ spl15_29 | ~ spl15_94), inference(avatar_contradiction_clause, [], [f1896])).
fof(f1896, plain, ($false | (~ spl15_29 | ~ spl15_94)), inference(subsumption_resolution, [], [f1895, f182])).
fof(f1895, plain, ((e0 = e2) | (~ spl15_29 | ~ spl15_94)), inference(forward_demodulation, [], [f1892, f433])).
fof(f1892, plain, ((e2 = op(e2, e0)) | (~ spl15_29 | ~ spl15_94)), inference(backward_demodulation, [], [f833, f433])).
fof(f1893, plain, (~ spl15_13 | ~ spl15_29), inference(avatar_split_clause, [], [f1886, f431, f363])).
fof(f1886, plain, (~ (e0 = op(e3, e0)) | ~ spl15_29), inference(backward_demodulation, [], [f138, f433])).
fof(f1885, plain, (~ spl15_38 | ~ spl15_40), inference(avatar_contradiction_clause, [], [f1884])).
fof(f1884, plain, ($false | (~ spl15_38 | ~ spl15_40)), inference(subsumption_resolution, [], [f1883, f185])).
fof(f1883, plain, ((e1 = e3) | (~ spl15_38 | ~ spl15_40)), inference(forward_demodulation, [], [f479, f471])).
fof(f1862, plain, (~ spl15_22 | ~ spl15_86 | spl15_106), inference(avatar_split_clause, [], [f1860, f887, f797, f401])).
fof(f1860, plain, (~ (e1 = op(e2, e2)) | (~ spl15_86 | spl15_106)), inference(backward_demodulation, [], [f889, f799])).
fof(f1846, plain, (~ spl15_15 | ~ spl15_60 | ~ spl15_62 | spl15_114), inference(avatar_split_clause, [], [f1831, f927, f571, f562, f371])).
fof(f1831, plain, (~ (e2 = op(e3, e0)) | (~ spl15_60 | ~ spl15_62 | spl15_114)), inference(forward_demodulation, [], [f1830, f564])).
fof(f1830, plain, (~ (e2 = op(op(e0, e1), e0)) | (~ spl15_62 | spl15_114)), inference(forward_demodulation, [], [f929, f573])).
fof(f1844, plain, (~ spl15_13 | ~ spl15_95), inference(avatar_contradiction_clause, [], [f1843])).
fof(f1843, plain, ($false | (~ spl15_13 | ~ spl15_95)), inference(subsumption_resolution, [], [f1842, f183])).
fof(f1842, plain, ((e0 = e3) | (~ spl15_13 | ~ spl15_95)), inference(forward_demodulation, [], [f1841, f365])).
fof(f1841, plain, ((e3 = op(e3, e0)) | (~ spl15_13 | ~ spl15_95)), inference(forward_demodulation, [], [f837, f365])).
fof(f1824, plain, (~ spl15_7 | ~ spl15_11), inference(avatar_split_clause, [], [f1819, f354, f337])).
fof(f1819, plain, (~ (e2 = op(e3, e2)) | ~ spl15_11), inference(backward_demodulation, [], [f178, f356])).
fof(f1814, plain, (spl15_34 | ~ spl15_48 | ~ spl15_93), inference(avatar_contradiction_clause, [], [f1813])).
fof(f1813, plain, ($false | (spl15_34 | ~ spl15_48 | ~ spl15_93)), inference(subsumption_resolution, [], [f1809, f453])).
fof(f1809, plain, ((e1 = op(e1, e3)) | (~ spl15_48 | ~ spl15_93)), inference(backward_demodulation, [], [f829, f513])).
fof(f1804, plain, (~ spl15_12 | ~ spl15_60), inference(avatar_split_clause, [], [f1799, f562, f358])).
fof(f1799, plain, (~ (e3 = op(e3, e1)) | ~ spl15_60), inference(backward_demodulation, [], [f141, f564])).
fof(f1789, plain, (~ spl15_60 | ~ spl15_62 | spl15_109), inference(avatar_split_clause, [], [f1708, f901, f571, f562])).
fof(f1708, plain, (~ (e3 = op(e0, e1)) | (~ spl15_62 | spl15_109)), inference(forward_demodulation, [], [f903, f573])).
fof(f1778, plain, (~ spl15_46 | ~ spl15_41 | spl15_89), inference(avatar_split_clause, [], [f1777, f810, f482, f503])).
fof(f1777, plain, (~ (e1 = op(e1, e0)) | (~ spl15_41 | spl15_89)), inference(forward_demodulation, [], [f811, f484])).
fof(f1761, plain, (~ spl15_13 | ~ spl15_16), inference(avatar_contradiction_clause, [], [f1760])).
fof(f1760, plain, ($false | (~ spl15_13 | ~ spl15_16)), inference(subsumption_resolution, [], [f1758, f183])).
fof(f1758, plain, ((e0 = e3) | (~ spl15_13 | ~ spl15_16)), inference(backward_demodulation, [], [f377, f365])).
fof(f1756, plain, (~ spl15_41 | ~ spl15_44), inference(avatar_contradiction_clause, [], [f1755])).
fof(f1755, plain, ($false | (~ spl15_41 | ~ spl15_44)), inference(subsumption_resolution, [], [f1754, f183])).
fof(f1754, plain, ((e0 = e3) | (~ spl15_41 | ~ spl15_44)), inference(backward_demodulation, [], [f496, f484])).
fof(f1751, plain, (spl15_42 | ~ spl15_46 | ~ spl15_93), inference(avatar_contradiction_clause, [], [f1750])).
fof(f1750, plain, ($false | (spl15_42 | ~ spl15_46 | ~ spl15_93)), inference(subsumption_resolution, [], [f1749, f487])).
fof(f1749, plain, ((e1 = op(e1, e1)) | (~ spl15_46 | ~ spl15_93)), inference(backward_demodulation, [], [f829, f505])).
fof(f1727, plain, (~ spl15_12 | ~ spl15_44), inference(avatar_split_clause, [], [f1726, f494, f358])).
fof(f1726, plain, (~ (e3 = op(e3, e1)) | ~ spl15_44), inference(forward_demodulation, [], [f143, f496])).
fof(f1725, plain, (~ spl15_12 | ~ spl15_2 | spl15_83), inference(avatar_split_clause, [], [f1567, f784, f316, f358])).
fof(f1567, plain, (~ (e3 = op(e3, e1)) | (~ spl15_2 | spl15_83)), inference(forward_demodulation, [], [f785, f318])).
fof(f785, plain, (~ (e3 = op(e3, op(e3, e3))) | spl15_83), inference(avatar_component_clause, [], [f784])).
fof(f1719, plain, (~ spl15_45 | ~ spl15_93), inference(avatar_contradiction_clause, [], [f1718])).
fof(f1718, plain, ($false | (~ spl15_45 | ~ spl15_93)), inference(subsumption_resolution, [], [f1717, f181])).
fof(f1717, plain, ((e0 = e1) | (~ spl15_45 | ~ spl15_93)), inference(forward_demodulation, [], [f1716, f501])).
fof(f1716, plain, ((e1 = op(e1, e0)) | (~ spl15_45 | ~ spl15_93)), inference(forward_demodulation, [], [f829, f501])).
fof(f1679, plain, (~ spl15_27 | ~ spl15_31), inference(avatar_split_clause, [], [f1674, f439, f422])).
fof(f1674, plain, (~ (e2 = op(e2, e1)) | ~ spl15_31), inference(backward_demodulation, [], [f169, f441])).
fof(f1678, plain, (~ spl15_15 | ~ spl15_31), inference(avatar_split_clause, [], [f1673, f439, f371])).
fof(f1673, plain, (~ (e2 = op(e3, e0)) | ~ spl15_31), inference(backward_demodulation, [], [f138, f441])).
fof(f1665, plain, (spl15_53 | ~ spl15_59 | ~ spl15_88), inference(avatar_contradiction_clause, [], [f1664])).
fof(f1664, plain, ($false | (spl15_53 | ~ spl15_59 | ~ spl15_88)), inference(subsumption_resolution, [], [f1659, f534])).
fof(f534, plain, (~ (e0 = op(e0, e2)) | spl15_53), inference(avatar_component_clause, [], [f533])).
fof(f1659, plain, ((e0 = op(e0, e2)) | (~ spl15_59 | ~ spl15_88)), inference(backward_demodulation, [], [f808, f560])).
fof(f808, plain, ((e0 = op(e0, op(e0, e1))) | ~ spl15_88), inference(avatar_component_clause, [], [f806])).
fof(f806, plain, (spl15_88 <=> (e0 = op(e0, op(e0, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl15_88])])).
fof(f1657, plain, (~ spl15_59 | ~ spl15_62 | spl15_111), inference(avatar_split_clause, [], [f1653, f911, f571, f558])).
fof(f1653, plain, (~ (e2 = op(e0, e1)) | (~ spl15_62 | spl15_111)), inference(backward_demodulation, [], [f913, f573])).
fof(f913, plain, (~ (e2 = op(e0, op(e0, e0))) | spl15_111), inference(avatar_component_clause, [], [f911])).
fof(f1655, plain, (~ spl15_30 | ~ spl15_62), inference(avatar_split_clause, [], [f1646, f571, f435])).
fof(f1646, plain, (~ (e1 = op(e2, e0)) | ~ spl15_62), inference(backward_demodulation, [], [f134, f573])).
fof(f1637, plain, (spl15_31 | ~ spl15_21 | ~ spl15_86), inference(avatar_split_clause, [], [f1628, f797, f397, f439])).
fof(f1628, plain, ((e2 = op(e2, e0)) | (~ spl15_21 | ~ spl15_86)), inference(forward_demodulation, [], [f799, f399])).
fof(f1625, plain, (~ spl15_9 | ~ spl15_15 | ~ spl15_91), inference(avatar_contradiction_clause, [], [f1624])).
fof(f1624, plain, ($false | (~ spl15_9 | ~ spl15_15 | ~ spl15_91)), inference(subsumption_resolution, [], [f1623, f186])).
fof(f1623, plain, ((e2 = e3) | (~ spl15_9 | ~ spl15_15 | ~ spl15_91)), inference(forward_demodulation, [], [f1622, f373])).
fof(f1622, plain, ((e3 = op(e3, e0)) | (~ spl15_9 | ~ spl15_91)), inference(forward_demodulation, [], [f820, f348])).
fof(f1620, plain, (spl15_53 | ~ spl15_55 | ~ spl15_84), inference(avatar_contradiction_clause, [], [f1619])).
fof(f1619, plain, ($false | (spl15_53 | ~ spl15_55 | ~ spl15_84)), inference(subsumption_resolution, [], [f1618, f534])).
fof(f1618, plain, ((e0 = op(e0, e2)) | (~ spl15_55 | ~ spl15_84)), inference(forward_demodulation, [], [f791, f543])).
fof(f1601, plain, (~ spl15_21 | ~ spl15_22), inference(avatar_contradiction_clause, [], [f1600])).
fof(f1600, plain, ($false | (~ spl15_21 | ~ spl15_22)), inference(subsumption_resolution, [], [f1599, f181])).
fof(f1599, plain, ((e0 = e1) | (~ spl15_21 | ~ spl15_22)), inference(backward_demodulation, [], [f403, f399])).
fof(f1596, plain, (~ spl15_26 | ~ spl15_30), inference(avatar_split_clause, [], [f1592, f435, f418])).
fof(f1592, plain, (~ (e1 = op(e2, e1)) | ~ spl15_30), inference(backward_demodulation, [], [f169, f437])).
fof(f1590, plain, (~ spl15_38 | spl15_42 | ~ spl15_85), inference(avatar_contradiction_clause, [], [f1589])).
fof(f1589, plain, ($false | (~ spl15_38 | spl15_42 | ~ spl15_85)), inference(subsumption_resolution, [], [f1587, f487])).
fof(f1587, plain, ((e1 = op(e1, e1)) | (~ spl15_38 | ~ spl15_85)), inference(backward_demodulation, [], [f795, f471])).
fof(f1576, plain, (spl15_57 | ~ spl15_58 | ~ spl15_88), inference(avatar_contradiction_clause, [], [f1575])).
fof(f1575, plain, ($false | (spl15_57 | ~ spl15_58 | ~ spl15_88)), inference(subsumption_resolution, [], [f1572, f551])).
fof(f1572, plain, ((e0 = op(e0, e1)) | (~ spl15_58 | ~ spl15_88)), inference(backward_demodulation, [], [f808, f556])).
fof(f1570, plain, (spl15_38 | ~ spl15_35 | ~ spl15_81), inference(avatar_split_clause, [], [f1569, f776, f456, f469])).
fof(f1569, plain, ((e1 = op(e1, e2)) | (~ spl15_35 | ~ spl15_81)), inference(forward_demodulation, [], [f778, f458])).
fof(f1541, plain, (~ spl15_9 | ~ spl15_11), inference(avatar_contradiction_clause, [], [f1540])).
fof(f1540, plain, ($false | (~ spl15_9 | ~ spl15_11)), inference(subsumption_resolution, [], [f1539, f182])).
fof(f1539, plain, ((e0 = e2) | (~ spl15_9 | ~ spl15_11)), inference(forward_demodulation, [], [f356, f348])).
fof(f1523, plain, (spl15_23 | ~ spl15_31 | ~ spl15_94), inference(avatar_contradiction_clause, [], [f1522])).
fof(f1522, plain, ($false | (spl15_23 | ~ spl15_31 | ~ spl15_94)), inference(subsumption_resolution, [], [f1519, f406])).
fof(f1519, plain, ((e2 = op(e2, e2)) | (~ spl15_31 | ~ spl15_94)), inference(backward_demodulation, [], [f833, f441])).
fof(f1497, plain, (~ spl15_62 | ~ spl15_64), inference(avatar_contradiction_clause, [], [f1496])).
fof(f1496, plain, ($false | (~ spl15_62 | ~ spl15_64)), inference(subsumption_resolution, [], [f1494, f185])).
fof(f1494, plain, ((e1 = e3) | (~ spl15_62 | ~ spl15_64)), inference(backward_demodulation, [], [f581, f573])).
fof(f1492, plain, (~ spl15_2 | ~ spl15_9 | ~ spl15_83), inference(avatar_contradiction_clause, [], [f1491])).
fof(f1491, plain, ($false | (~ spl15_2 | ~ spl15_9 | ~ spl15_83)), inference(subsumption_resolution, [], [f1490, f183])).
fof(f1490, plain, ((e0 = e3) | (~ spl15_2 | ~ spl15_9 | ~ spl15_83)), inference(forward_demodulation, [], [f1489, f348])).
fof(f1489, plain, ((e3 = op(e3, e1)) | (~ spl15_2 | ~ spl15_83)), inference(forward_demodulation, [], [f786, f318])).
fof(f1463, plain, (~ spl15_35 | ~ spl15_44 | spl15_104), inference(avatar_contradiction_clause, [], [f1462])).
fof(f1462, plain, ($false | (~ spl15_35 | ~ spl15_44 | spl15_104)), inference(subsumption_resolution, [], [f1461, f458])).
fof(f1461, plain, (~ (e2 = op(e1, e3)) | (~ spl15_44 | spl15_104)), inference(forward_demodulation, [], [f879, f496])).
fof(f1450, plain, (spl15_34 | ~ spl15_44 | ~ spl15_89), inference(avatar_contradiction_clause, [], [f1449])).
fof(f1449, plain, ($false | (spl15_34 | ~ spl15_44 | ~ spl15_89)), inference(subsumption_resolution, [], [f1448, f453])).
fof(f1448, plain, ((e1 = op(e1, e3)) | (~ spl15_44 | ~ spl15_89)), inference(forward_demodulation, [], [f812, f496])).
fof(f1439, plain, (~ spl15_42 | ~ spl15_44), inference(avatar_contradiction_clause, [], [f1438])).
fof(f1438, plain, ($false | (~ spl15_42 | ~ spl15_44)), inference(subsumption_resolution, [], [f1437, f185])).
fof(f1437, plain, ((e1 = e3) | (~ spl15_42 | ~ spl15_44)), inference(forward_demodulation, [], [f496, f488])).
fof(f1432, plain, (~ spl15_49 | spl15_61 | ~ spl15_80), inference(avatar_contradiction_clause, [], [f1431])).
fof(f1431, plain, ($false | (~ spl15_49 | spl15_61 | ~ spl15_80)), inference(subsumption_resolution, [], [f1428, f568])).
fof(f568, plain, (~ (e0 = op(e0, e0)) | spl15_61), inference(avatar_component_clause, [], [f567])).
fof(f1428, plain, ((e0 = op(e0, e0)) | (~ spl15_49 | ~ spl15_80)), inference(backward_demodulation, [], [f774, f518])).
fof(f1420, plain, (~ spl15_60 | ~ spl15_64), inference(avatar_split_clause, [], [f1416, f579, f562])).
fof(f1416, plain, (~ (e3 = op(e0, e1)) | ~ spl15_64), inference(backward_demodulation, [], [f157, f581])).
fof(f1399, plain, (~ spl15_33 | ~ spl15_35), inference(avatar_contradiction_clause, [], [f1398])).
fof(f1398, plain, ($false | (~ spl15_33 | ~ spl15_35)), inference(subsumption_resolution, [], [f1397, f182])).
fof(f1397, plain, ((e0 = e2) | (~ spl15_33 | ~ spl15_35)), inference(forward_demodulation, [], [f458, f450])).
fof(f1396, plain, (~ spl15_33 | ~ spl15_40 | ~ spl15_85), inference(avatar_contradiction_clause, [], [f1395])).
fof(f1395, plain, ($false | (~ spl15_33 | ~ spl15_40 | ~ spl15_85)), inference(subsumption_resolution, [], [f1394, f181])).
fof(f1394, plain, ((e0 = e1) | (~ spl15_33 | ~ spl15_40 | ~ spl15_85)), inference(forward_demodulation, [], [f1393, f450])).
fof(f1393, plain, ((e1 = op(e1, e3)) | (~ spl15_40 | ~ spl15_85)), inference(backward_demodulation, [], [f795, f479])).
fof(f1392, plain, (spl15_49 | ~ spl15_52 | ~ spl15_80), inference(avatar_contradiction_clause, [], [f1391])).
fof(f1391, plain, ($false | (spl15_49 | ~ spl15_52 | ~ spl15_80)), inference(subsumption_resolution, [], [f1390, f517])).
fof(f517, plain, (~ (e0 = op(e0, e3)) | spl15_49), inference(avatar_component_clause, [], [f516])).
fof(f1390, plain, ((e0 = op(e0, e3)) | (~ spl15_52 | ~ spl15_80)), inference(backward_demodulation, [], [f774, f530])).
fof(f1389, plain, (~ spl15_54 | ~ spl15_55), inference(avatar_contradiction_clause, [], [f1388])).
fof(f1388, plain, ($false | (~ spl15_54 | ~ spl15_55)), inference(subsumption_resolution, [], [f1387, f184])).
fof(f1387, plain, ((e1 = e2) | (~ spl15_54 | ~ spl15_55)), inference(forward_demodulation, [], [f543, f539])).
fof(f1379, plain, (~ spl15_52 | ~ spl15_60), inference(avatar_split_clause, [], [f1378, f562, f528])).
fof(f1378, plain, (~ (e3 = op(e0, e3)) | ~ spl15_60), inference(forward_demodulation, [], [f161, f564])).
fof(f1377, plain, (~ spl15_40 | ~ spl15_8), inference(avatar_split_clause, [], [f1376, f341, f477])).
fof(f1376, plain, (~ (e3 = op(e1, e2)) | ~ spl15_8), inference(forward_demodulation, [], [f149, f343])).
fof(f1373, plain, (~ spl15_40 | ~ spl15_48), inference(avatar_split_clause, [], [f1372, f511, f477])).
fof(f1372, plain, (~ (e3 = op(e1, e2)) | ~ spl15_48), inference(forward_demodulation, [], [f164, f513])).
fof(f1371, plain, (~ spl15_40 | ~ spl15_103 | spl15_117), inference(avatar_split_clause, [], [f1344, f942, f872, f477])).
fof(f1344, plain, (~ (e3 = op(e1, e2)) | (~ spl15_103 | spl15_117)), inference(backward_demodulation, [], [f944, f873])).
fof(f1356, plain, (spl15_19 | ~ spl15_20 | ~ spl15_82), inference(avatar_contradiction_clause, [], [f1355])).
fof(f1355, plain, ($false | (spl15_19 | ~ spl15_20 | ~ spl15_82)), inference(subsumption_resolution, [], [f1354, f389])).
fof(f389, plain, (~ (e2 = op(e2, e3)) | spl15_19), inference(avatar_component_clause, [], [f388])).
fof(f1354, plain, ((e2 = op(e2, e3)) | (~ spl15_20 | ~ spl15_82)), inference(forward_demodulation, [], [f782, f394])).
fof(f394, plain, ((e3 = op(e2, e3)) | ~ spl15_20), inference(avatar_component_clause, [], [f392])).
fof(f1352, plain, (spl15_38 | ~ spl15_39 | ~ spl15_85), inference(avatar_contradiction_clause, [], [f1351])).
fof(f1351, plain, ($false | (spl15_38 | ~ spl15_39 | ~ spl15_85)), inference(subsumption_resolution, [], [f1350, f470])).
fof(f470, plain, (~ (e1 = op(e1, e2)) | spl15_38), inference(avatar_component_clause, [], [f469])).
fof(f1350, plain, ((e1 = op(e1, e2)) | (~ spl15_39 | ~ spl15_85)), inference(forward_demodulation, [], [f795, f475])).
fof(f1349, plain, (~ spl15_49 | ~ spl15_60 | spl15_88), inference(avatar_split_clause, [], [f1348, f806, f562, f516])).
fof(f1348, plain, (~ (e0 = op(e0, e3)) | (~ spl15_60 | spl15_88)), inference(forward_demodulation, [], [f807, f564])).
fof(f807, plain, (~ (e0 = op(e0, op(e0, e1))) | spl15_88), inference(avatar_component_clause, [], [f806])).
fof(f1328, plain, (~ spl15_5 | ~ spl15_8), inference(avatar_contradiction_clause, [], [f1327])).
fof(f1327, plain, ($false | (~ spl15_5 | ~ spl15_8)), inference(subsumption_resolution, [], [f1325, f183])).
fof(f1325, plain, ((e0 = e3) | (~ spl15_5 | ~ spl15_8)), inference(backward_demodulation, [], [f343, f331])).
fof(f1314, plain, (~ spl15_18 | ~ spl15_24 | spl15_103), inference(avatar_split_clause, [], [f1312, f872, f409, f384])).
fof(f1312, plain, (~ (e1 = op(e2, e3)) | (~ spl15_24 | spl15_103)), inference(backward_demodulation, [], [f874, f411])).
fof(f874, plain, (~ (e1 = op(e2, op(e2, e2))) | spl15_103), inference(avatar_component_clause, [], [f872])).
fof(f1305, plain, (~ spl15_30 | ~ spl15_31), inference(avatar_contradiction_clause, [], [f1304])).
fof(f1304, plain, ($false | (~ spl15_30 | ~ spl15_31)), inference(subsumption_resolution, [], [f1303, f184])).
fof(f1303, plain, ((e1 = e2) | (~ spl15_30 | ~ spl15_31)), inference(backward_demodulation, [], [f441, f437])).
fof(f1297, plain, (~ spl15_49 | ~ spl15_51), inference(avatar_contradiction_clause, [], [f1296])).
fof(f1296, plain, ($false | (~ spl15_49 | ~ spl15_51)), inference(subsumption_resolution, [], [f1295, f182])).
fof(f1295, plain, ((e0 = e2) | (~ spl15_49 | ~ spl15_51)), inference(backward_demodulation, [], [f526, f518])).
fof(f1292, plain, (~ spl15_59 | ~ spl15_60), inference(avatar_contradiction_clause, [], [f1291])).
fof(f1291, plain, ($false | (~ spl15_59 | ~ spl15_60)), inference(subsumption_resolution, [], [f1290, f186])).
fof(f1290, plain, ((e2 = e3) | (~ spl15_59 | ~ spl15_60)), inference(backward_demodulation, [], [f564, f560])).
fof(f1278, plain, (~ spl15_62 | ~ spl15_54), inference(avatar_split_clause, [], [f1277, f537, f571])).
fof(f1277, plain, (~ (op(e0, e0) = e1) | ~ spl15_54), inference(forward_demodulation, [], [f158, f539])).
fof(f1276, plain, (~ spl15_58 | ~ spl15_42), inference(avatar_split_clause, [], [f1275, f486, f554])).
fof(f1275, plain, (~ (e1 = op(e0, e1)) | ~ spl15_42), inference(forward_demodulation, [], [f139, f488])).
fof(f1271, plain, (~ spl15_58 | ~ spl15_54), inference(avatar_split_clause, [], [f1270, f537, f554])).
fof(f1270, plain, (~ (e1 = op(e0, e1)) | ~ spl15_54), inference(forward_demodulation, [], [f160, f539])).
fof(f1269, plain, (~ spl15_59 | ~ spl15_51), inference(avatar_split_clause, [], [f1268, f524, f558])).
fof(f1268, plain, (~ (e2 = op(e0, e1)) | ~ spl15_51), inference(forward_demodulation, [], [f161, f526])).
fof(f1267, plain, (spl15_49 | ~ spl15_60 | ~ spl15_88), inference(avatar_split_clause, [], [f1266, f806, f562, f516])).
fof(f1266, plain, ((e0 = op(e0, e3)) | (~ spl15_60 | ~ spl15_88)), inference(forward_demodulation, [], [f808, f564])).
fof(f1263, plain, (spl15_57 | ~ spl15_54 | ~ spl15_84), inference(avatar_split_clause, [], [f1192, f789, f537, f550])).
fof(f1192, plain, ((e0 = op(e0, e1)) | (~ spl15_54 | ~ spl15_84)), inference(backward_demodulation, [], [f791, f539])).
fof(f1256, plain, (~ spl15_47 | ~ spl15_39), inference(avatar_split_clause, [], [f1255, f473, f507])).
fof(f1255, plain, (~ (e2 = op(e1, e0)) | ~ spl15_39), inference(forward_demodulation, [], [f164, f475])).
fof(f1251, plain, (~ spl15_34 | ~ spl15_2), inference(avatar_split_clause, [], [f1150, f316, f452])).
fof(f1150, plain, (~ (e1 = op(e1, e3)) | ~ spl15_2), inference(forward_demodulation, [], [f155, f318])).
fof(f1249, plain, (~ spl15_34 | ~ spl15_42), inference(avatar_split_clause, [], [f1248, f486, f452])).
fof(f1248, plain, (~ (e1 = op(e1, e3)) | ~ spl15_42), inference(forward_demodulation, [], [f167, f488])).
fof(f1236, plain, (~ spl15_28 | ~ spl15_20), inference(avatar_split_clause, [], [f1234, f392, f426])).
fof(f1234, plain, (~ (e3 = op(e2, e1)) | ~ spl15_20), inference(backward_demodulation, [], [f173, f394])).
fof(f1221, plain, (~ spl15_1 | ~ spl15_2), inference(avatar_contradiction_clause, [], [f1220])).
fof(f1220, plain, ($false | (~ spl15_1 | ~ spl15_2)), inference(subsumption_resolution, [], [f1219, f181])).
fof(f1219, plain, ((e0 = e1) | (~ spl15_1 | ~ spl15_2)), inference(backward_demodulation, [], [f318, f314])).
fof(f1193, plain, (~ spl15_38 | ~ spl15_54), inference(avatar_split_clause, [], [f1191, f537, f469])).
fof(f1191, plain, (~ (e1 = op(e1, e2)) | ~ spl15_54), inference(backward_demodulation, [], [f145, f539])).
fof(f1190, plain, (~ spl15_57 | ~ spl15_60), inference(avatar_contradiction_clause, [], [f1189])).
fof(f1189, plain, ($false | (~ spl15_57 | ~ spl15_60)), inference(subsumption_resolution, [], [f1188, f183])).
fof(f1188, plain, ((e0 = e3) | (~ spl15_57 | ~ spl15_60)), inference(forward_demodulation, [], [f564, f552])).
fof(f1187, plain, (~ spl15_53 | ~ spl15_61), inference(avatar_split_clause, [], [f1182, f567, f533])).
fof(f1182, plain, (~ (e0 = op(e0, e2)) | ~ spl15_61), inference(backward_demodulation, [], [f158, f569])).
fof(f1186, plain, (~ spl15_29 | ~ spl15_61), inference(avatar_split_clause, [], [f1181, f567, f431])).
fof(f1181, plain, (~ (e0 = op(e2, e0)) | ~ spl15_61), inference(backward_demodulation, [], [f134, f569])).
fof(f1177, plain, (~ spl15_61 | ~ spl15_45), inference(avatar_split_clause, [], [f1176, f499, f567])).
fof(f1176, plain, (~ (e0 = op(e0, e0)) | ~ spl15_45), inference(forward_demodulation, [], [f133, f501])).
fof(f1174, plain, (~ spl15_61 | ~ spl15_57), inference(avatar_split_clause, [], [f1173, f550, f567])).
fof(f1173, plain, (~ (e0 = op(e0, e0)) | ~ spl15_57), inference(forward_demodulation, [], [f157, f552])).
fof(f1165, plain, (spl15_53 | ~ spl15_51 | ~ spl15_80), inference(avatar_split_clause, [], [f1001, f772, f524, f533])).
fof(f1001, plain, ((e0 = op(e0, e2)) | (~ spl15_51 | ~ spl15_80)), inference(backward_demodulation, [], [f774, f526])).
fof(f1157, plain, (~ spl15_38 | ~ spl15_34), inference(avatar_split_clause, [], [f1156, f452, f469])).
fof(f1156, plain, (~ (e1 = op(e1, e2)) | ~ spl15_34), inference(forward_demodulation, [], [f168, f454])).
fof(f1145, plain, (~ spl15_30 | ~ spl15_18), inference(avatar_split_clause, [], [f1144, f384, f435])).
fof(f1144, plain, (~ (e1 = op(e2, e0)) | ~ spl15_18), inference(forward_demodulation, [], [f171, f386])).
fof(f1139, plain, (~ spl15_26 | ~ spl15_18), inference(avatar_split_clause, [], [f1138, f384, f418])).
fof(f1138, plain, (~ (e1 = op(e2, e1)) | ~ spl15_18), inference(forward_demodulation, [], [f173, f386])).
fof(f1137, plain, (~ spl15_18 | ~ spl15_20), inference(avatar_contradiction_clause, [], [f1136])).
fof(f1136, plain, ($false | (~ spl15_18 | ~ spl15_20)), inference(subsumption_resolution, [], [f1135, f185])).
fof(f1135, plain, ((e1 = e3) | (~ spl15_18 | ~ spl15_20)), inference(forward_demodulation, [], [f394, f386])).
fof(f1122, plain, (~ spl15_11 | ~ spl15_2 | spl15_97), inference(avatar_split_clause, [], [f1115, f844, f316, f354])).
fof(f1115, plain, (~ (e2 = op(e3, e1)) | (~ spl15_2 | spl15_97)), inference(backward_demodulation, [], [f846, f318])).
fof(f1121, plain, (~ spl15_6 | ~ spl15_8), inference(avatar_contradiction_clause, [], [f1120])).
fof(f1120, plain, ($false | (~ spl15_6 | ~ spl15_8)), inference(subsumption_resolution, [], [f1119, f185])).
fof(f1119, plain, ((e1 = e3) | (~ spl15_6 | ~ spl15_8)), inference(forward_demodulation, [], [f343, f335])).
fof(f1113, plain, (~ spl15_5 | ~ spl15_6), inference(avatar_contradiction_clause, [], [f1112])).
fof(f1112, plain, ($false | (~ spl15_5 | ~ spl15_6)), inference(subsumption_resolution, [], [f1111, f181])).
fof(f1111, plain, ((e0 = e1) | (~ spl15_5 | ~ spl15_6)), inference(backward_demodulation, [], [f335, f331])).
fof(f1110, plain, (~ spl15_6 | ~ spl15_7), inference(avatar_contradiction_clause, [], [f1109])).
fof(f1109, plain, ($false | (~ spl15_6 | ~ spl15_7)), inference(subsumption_resolution, [], [f1108, f184])).
fof(f1108, plain, ((e1 = e2) | (~ spl15_6 | ~ spl15_7)), inference(backward_demodulation, [], [f339, f335])).
fof(f1107, plain, (~ spl15_7 | ~ spl15_8), inference(avatar_contradiction_clause, [], [f1106])).
fof(f1106, plain, ($false | (~ spl15_7 | ~ spl15_8)), inference(subsumption_resolution, [], [f1105, f186])).
fof(f1105, plain, ((e2 = e3) | (~ spl15_7 | ~ spl15_8)), inference(backward_demodulation, [], [f343, f339])).
fof(f1102, plain, (~ spl15_10 | ~ spl15_11), inference(avatar_contradiction_clause, [], [f1101])).
fof(f1101, plain, ($false | (~ spl15_10 | ~ spl15_11)), inference(subsumption_resolution, [], [f1100, f184])).
fof(f1100, plain, ((e1 = e2) | (~ spl15_10 | ~ spl15_11)), inference(backward_demodulation, [], [f356, f352])).
fof(f1099, plain, (~ spl15_11 | ~ spl15_12), inference(avatar_contradiction_clause, [], [f1098])).
fof(f1098, plain, ($false | (~ spl15_11 | ~ spl15_12)), inference(subsumption_resolution, [], [f1097, f186])).
fof(f1097, plain, ((e2 = e3) | (~ spl15_11 | ~ spl15_12)), inference(backward_demodulation, [], [f360, f356])).
fof(f1096, plain, (~ spl15_4 | ~ spl15_12), inference(avatar_split_clause, [], [f1094, f358, f324])).
fof(f1094, plain, (~ (e3 = op(e3, e3)) | ~ spl15_12), inference(backward_demodulation, [], [f179, f360])).
fof(f1092, plain, (~ spl15_14 | ~ spl15_15), inference(avatar_contradiction_clause, [], [f1091])).
fof(f1091, plain, ($false | (~ spl15_14 | ~ spl15_15)), inference(subsumption_resolution, [], [f1090, f184])).
fof(f1090, plain, ((e1 = e2) | (~ spl15_14 | ~ spl15_15)), inference(backward_demodulation, [], [f373, f369])).
fof(f1083, plain, (~ spl15_17 | ~ spl15_18), inference(avatar_contradiction_clause, [], [f1082])).
fof(f1082, plain, ($false | (~ spl15_17 | ~ spl15_18)), inference(subsumption_resolution, [], [f1081, f181])).
fof(f1081, plain, ((e0 = e1) | (~ spl15_17 | ~ spl15_18)), inference(backward_demodulation, [], [f386, f382])).
fof(f1080, plain, (~ spl15_18 | ~ spl15_19), inference(avatar_contradiction_clause, [], [f1079])).
fof(f1079, plain, ($false | (~ spl15_18 | ~ spl15_19)), inference(subsumption_resolution, [], [f1078, f184])).
fof(f1078, plain, ((e1 = e2) | (~ spl15_18 | ~ spl15_19)), inference(backward_demodulation, [], [f390, f386])).
fof(f1077, plain, (~ spl15_19 | ~ spl15_20), inference(avatar_contradiction_clause, [], [f1076])).
fof(f1076, plain, ($false | (~ spl15_19 | ~ spl15_20)), inference(subsumption_resolution, [], [f1075, f186])).
fof(f1075, plain, ((e2 = e3) | (~ spl15_19 | ~ spl15_20)), inference(backward_demodulation, [], [f394, f390])).
fof(f1074, plain, (~ spl15_4 | ~ spl15_20), inference(avatar_split_clause, [], [f1073, f392, f324])).
fof(f1073, plain, (~ (e3 = op(e3, e3)) | ~ spl15_20), inference(backward_demodulation, [], [f156, f394])).
fof(f1072, plain, (~ spl15_6 | ~ spl15_21 | ~ spl15_32 | spl15_106), inference(avatar_split_clause, [], [f1071, f887, f443, f397, f333])).
fof(f1071, plain, (~ (e1 = op(e3, e2)) | (~ spl15_21 | ~ spl15_32 | spl15_106)), inference(forward_demodulation, [], [f1067, f445])).
fof(f1067, plain, (~ (e1 = op(op(e2, e0), e2)) | (~ spl15_21 | spl15_106)), inference(backward_demodulation, [], [f889, f399])).
fof(f1070, plain, (~ spl15_30 | ~ spl15_21 | spl15_103), inference(avatar_split_clause, [], [f1066, f872, f397, f435])).
fof(f1066, plain, (~ (e1 = op(e2, e0)) | (~ spl15_21 | spl15_103)), inference(backward_demodulation, [], [f874, f399])).
fof(f1069, plain, (~ spl15_17 | ~ spl15_21), inference(avatar_split_clause, [], [f1063, f397, f380])).
fof(f1063, plain, (~ (e0 = op(e2, e3)) | ~ spl15_21), inference(backward_demodulation, [], [f174, f399])).
fof(f1061, plain, (~ spl15_17 | ~ spl15_25), inference(avatar_split_clause, [], [f1058, f414, f380])).
fof(f1058, plain, (~ (e0 = op(e2, e3)) | ~ spl15_25), inference(backward_demodulation, [], [f173, f416])).
fof(f1047, plain, (~ spl15_34 | ~ spl15_35), inference(avatar_contradiction_clause, [], [f1046])).
fof(f1046, plain, ($false | (~ spl15_34 | ~ spl15_35)), inference(subsumption_resolution, [], [f1045, f184])).
fof(f1045, plain, ((e1 = e2) | (~ spl15_34 | ~ spl15_35)), inference(backward_demodulation, [], [f458, f454])).
fof(f1044, plain, (~ spl15_35 | ~ spl15_36), inference(avatar_contradiction_clause, [], [f1043])).
fof(f1043, plain, ($false | (~ spl15_35 | ~ spl15_36)), inference(subsumption_resolution, [], [f1042, f186])).
fof(f1042, plain, ((e2 = e3) | (~ spl15_35 | ~ spl15_36)), inference(backward_demodulation, [], [f462, f458])).
fof(f1028, plain, (~ spl15_33 | ~ spl15_41), inference(avatar_split_clause, [], [f1021, f482, f448])).
fof(f1021, plain, (~ (e0 = op(e1, e3)) | ~ spl15_41), inference(backward_demodulation, [], [f167, f484])).
fof(f1015, plain, (~ spl15_41 | ~ spl15_45), inference(avatar_split_clause, [], [f1010, f499, f482])).
fof(f1010, plain, (~ (e0 = op(e1, e1)) | ~ spl15_45), inference(backward_demodulation, [], [f163, f501])).
fof(f1007, plain, (~ spl15_50 | ~ spl15_51), inference(avatar_contradiction_clause, [], [f1006])).
fof(f1006, plain, ($false | (~ spl15_50 | ~ spl15_51)), inference(subsumption_resolution, [], [f1005, f184])).
fof(f1005, plain, ((e1 = e2) | (~ spl15_50 | ~ spl15_51)), inference(backward_demodulation, [], [f526, f522])).
fof(f1004, plain, (~ spl15_3 | ~ spl15_51), inference(avatar_split_clause, [], [f1000, f524, f320])).
fof(f1000, plain, (~ (e2 = op(e3, e3)) | ~ spl15_51), inference(backward_demodulation, [], [f153, f526])).
fof(f997, plain, (spl15_61 | ~ spl15_53 | ~ spl15_84), inference(avatar_split_clause, [], [f992, f789, f533, f567])).
fof(f992, plain, ((e0 = op(e0, e0)) | (~ spl15_53 | ~ spl15_84)), inference(backward_demodulation, [], [f791, f535])).
fof(f996, plain, (~ spl15_49 | ~ spl15_53), inference(avatar_split_clause, [], [f991, f533, f516])).
fof(f991, plain, (~ (e0 = op(e0, e3)) | ~ spl15_53), inference(backward_demodulation, [], [f162, f535])).
fof(f987, plain, (spl15_61 | ~ spl15_57 | ~ spl15_88), inference(avatar_split_clause, [], [f981, f806, f550, f567])).
fof(f981, plain, ((e0 = op(e0, e0)) | (~ spl15_57 | ~ spl15_88)), inference(backward_demodulation, [], [f808, f552])).
fof(f985, plain, (~ spl15_53 | ~ spl15_57), inference(avatar_split_clause, [], [f979, f550, f533])).
fof(f979, plain, (~ (e0 = op(e0, e2)) | ~ spl15_57), inference(backward_demodulation, [], [f160, f552])).
fof(f984, plain, (~ spl15_9 | ~ spl15_57), inference(avatar_split_clause, [], [f978, f550, f346])).
fof(f978, plain, (~ (e0 = op(e3, e1)) | ~ spl15_57), inference(backward_demodulation, [], [f141, f552])).
fof(f957, plain, (~ spl15_118 | ~ spl15_43 | ~ spl15_115), inference(avatar_split_clause, [], [f309, f932, f490, f947])).
fof(f309, plain, (~ (e0 = op(e1, op(e1, e1))) | ~ (e2 = op(e1, e1)) | ~ (e3 = op(op(e1, op(e1, e1)), e1))), inference(cnf_transformation, [], [f52])).
fof(f52, plain, (~ (e0 = op(e1, op(e1, e1))) | ~ (e2 = op(e1, e1)) | ~ (e3 = op(op(e1, op(e1, e1)), e1))), inference(ennf_transformation, [], [f28])).
fof(f28, plain, ~ ((e0 = op(e1, op(e1, e1))) & (e2 = op(e1, e1)) & (e3 = op(op(e1, op(e1, e1)), e1))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG131+1.p', ax28)).
fof(f956, plain, (~ spl15_119 | ~ spl15_62 | ~ spl15_111), inference(avatar_split_clause, [], [f308, f911, f571, f953])).
fof(f308, plain, (~ (e2 = op(e0, op(e0, e0))) | ~ (op(e0, e0) = e1) | ~ (e3 = op(op(e0, op(e0, e0)), e0))), inference(cnf_transformation, [], [f51])).
fof(f51, plain, (~ (e2 = op(e0, op(e0, e0))) | ~ (op(e0, e0) = e1) | ~ (e3 = op(op(e0, op(e0, e0)), e0))), inference(ennf_transformation, [], [f27])).
fof(f27, plain, ~ ((e2 = op(e0, op(e0, e0))) & (op(e0, e0) = e1) & (e3 = op(op(e0, op(e0, e0)), e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG131+1.p', ax27)).
fof(f951, plain, (~ spl15_117 | ~ spl15_22 | ~ spl15_110), inference(avatar_split_clause, [], [f307, f906, f401, f942])).
fof(f307, plain, (~ (e0 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, e2)) | ~ (e3 = op(op(e2, op(e2, e2)), e2))), inference(cnf_transformation, [], [f50])).
fof(f50, plain, (~ (e0 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, e2)) | ~ (e3 = op(op(e2, op(e2, e2)), e2))), inference(ennf_transformation, [], [f26])).
fof(f26, plain, ~ ((e0 = op(e2, op(e2, e2))) & (e1 = op(e2, e2)) & (e3 = op(op(e2, op(e2, e2)), e2))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG131+1.p', ax26)).
fof(f950, plain, (~ spl15_118 | ~ spl15_41 | ~ spl15_104), inference(avatar_split_clause, [], [f306, f877, f482, f947])).
fof(f306, plain, (~ (e2 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, e1)) | ~ (e3 = op(op(e1, op(e1, e1)), e1))), inference(cnf_transformation, [], [f49])).
fof(f49, plain, (~ (e2 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, e1)) | ~ (e3 = op(op(e1, op(e1, e1)), e1))), inference(ennf_transformation, [], [f25])).
fof(f25, plain, ~ ((e2 = op(e1, op(e1, e1))) & (e0 = op(e1, e1)) & (e3 = op(op(e1, op(e1, e1)), e1))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG131+1.p', ax25)).
fof(f945, plain, (~ spl15_117 | ~ spl15_21 | ~ spl15_103), inference(avatar_split_clause, [], [f305, f872, f397, f942])).
fof(f305, plain, (~ (e1 = op(e2, op(e2, e2))) | ~ (e0 = op(e2, e2)) | ~ (e3 = op(op(e2, op(e2, e2)), e2))), inference(cnf_transformation, [], [f48])).
fof(f48, plain, (~ (e1 = op(e2, op(e2, e2))) | ~ (e0 = op(e2, e2)) | ~ (e3 = op(op(e2, op(e2, e2)), e2))), inference(ennf_transformation, [], [f24])).
fof(f24, plain, ~ ((e1 = op(e2, op(e2, e2))) & (e0 = op(e2, e2)) & (e3 = op(op(e2, op(e2, e2)), e2))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG131+1.p', ax24)).
fof(f940, plain, (~ spl15_114 | ~ spl15_64 | ~ spl15_116), inference(avatar_split_clause, [], [f304, f937, f579, f927])).
fof(f304, plain, (~ (e1 = op(e0, op(e0, e0))) | ~ (op(e0, e0) = e3) | ~ (e2 = op(op(e0, op(e0, e0)), e0))), inference(cnf_transformation, [], [f47])).
fof(f47, plain, (~ (e1 = op(e0, op(e0, e0))) | ~ (op(e0, e0) = e3) | ~ (e2 = op(op(e0, op(e0, e0)), e0))), inference(ennf_transformation, [], [f23])).
fof(f23, plain, ~ ((e1 = op(e0, op(e0, e0))) & (op(e0, e0) = e3) & (e2 = op(op(e0, op(e0, e0)), e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG131+1.p', ax23)).
fof(f935, plain, (~ spl15_113 | ~ spl15_44 | ~ spl15_115), inference(avatar_split_clause, [], [f303, f932, f494, f921])).
fof(f303, plain, (~ (e0 = op(e1, op(e1, e1))) | ~ (e3 = op(e1, e1)) | ~ (e2 = op(op(e1, op(e1, e1)), e1))), inference(cnf_transformation, [], [f46])).
fof(f46, plain, (~ (e0 = op(e1, op(e1, e1))) | ~ (e3 = op(e1, e1)) | ~ (e2 = op(op(e1, op(e1, e1)), e1))), inference(ennf_transformation, [], [f22])).
fof(f22, plain, ~ ((e0 = op(e1, op(e1, e1))) & (e3 = op(e1, e1)) & (e2 = op(op(e1, op(e1, e1)), e1))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG131+1.p', ax22)).
fof(f930, plain, (~ spl15_114 | ~ spl15_62 | ~ spl15_109), inference(avatar_split_clause, [], [f302, f901, f571, f927])).
fof(f302, plain, (~ (e3 = op(e0, op(e0, e0))) | ~ (op(e0, e0) = e1) | ~ (e2 = op(op(e0, op(e0, e0)), e0))), inference(cnf_transformation, [], [f45])).
fof(f45, plain, (~ (e3 = op(e0, op(e0, e0))) | ~ (op(e0, e0) = e1) | ~ (e2 = op(op(e0, op(e0, e0)), e0))), inference(ennf_transformation, [], [f21])).
fof(f21, plain, ~ ((e3 = op(e0, op(e0, e0))) & (op(e0, e0) = e1) & (e2 = op(op(e0, op(e0, e0)), e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG131+1.p', ax21)).
fof(f919, plain, (~ spl15_112 | ~ spl15_1 | ~ spl15_100), inference(avatar_split_clause, [], [f299, f858, f312, f916])).
fof(f299, plain, (~ (e1 = op(e3, op(e3, e3))) | ~ (e0 = op(e3, e3)) | ~ (e2 = op(op(e3, op(e3, e3)), e3))), inference(cnf_transformation, [], [f42])).
fof(f42, plain, (~ (e1 = op(e3, op(e3, e3))) | ~ (e0 = op(e3, e3)) | ~ (e2 = op(op(e3, op(e3, e3)), e3))), inference(ennf_transformation, [], [f18])).
fof(f18, plain, ~ ((e1 = op(e3, op(e3, e3))) & (e0 = op(e3, e3)) & (e2 = op(op(e3, op(e3, e3)), e3))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG131+1.p', ax18)).
fof(f909, plain, (~ spl15_106 | ~ spl15_24 | ~ spl15_110), inference(avatar_split_clause, [], [f297, f906, f409, f887])).
fof(f297, plain, (~ (e0 = op(e2, op(e2, e2))) | ~ (e3 = op(e2, e2)) | ~ (e1 = op(op(e2, op(e2, e2)), e2))), inference(cnf_transformation, [], [f40])).
fof(f40, plain, (~ (e0 = op(e2, op(e2, e2))) | ~ (e3 = op(e2, e2)) | ~ (e1 = op(op(e2, op(e2, e2)), e2))), inference(ennf_transformation, [], [f16])).
fof(f16, plain, ~ ((e0 = op(e2, op(e2, e2))) & (e3 = op(e2, e2)) & (e1 = op(op(e2, op(e2, e2)), e2))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG131+1.p', ax16)).
fof(f904, plain, (~ spl15_108 | ~ spl15_63 | ~ spl15_109), inference(avatar_split_clause, [], [f296, f901, f575, f897])).
fof(f296, plain, (~ (e3 = op(e0, op(e0, e0))) | ~ (op(e0, e0) = e2) | ~ (e1 = op(op(e0, op(e0, e0)), e0))), inference(cnf_transformation, [], [f39])).
fof(f39, plain, (~ (e3 = op(e0, op(e0, e0))) | ~ (op(e0, e0) = e2) | ~ (e1 = op(op(e0, op(e0, e0)), e0))), inference(ennf_transformation, [], [f15])).
fof(f15, plain, ~ ((e3 = op(e0, op(e0, e0))) & (op(e0, e0) = e2) & (e1 = op(op(e0, op(e0, e0)), e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG131+1.p', ax15)).
fof(f895, plain, (~ spl15_105 | ~ spl15_3 | ~ spl15_107), inference(avatar_split_clause, [], [f295, f892, f320, f882])).
fof(f295, plain, (~ (e0 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, e3)) | ~ (e1 = op(op(e3, op(e3, e3)), e3))), inference(cnf_transformation, [], [f38])).
fof(f38, plain, (~ (e0 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, e3)) | ~ (e1 = op(op(e3, op(e3, e3)), e3))), inference(ennf_transformation, [], [f14])).
fof(f14, plain, ~ ((e0 = op(e3, op(e3, e3))) & (e2 = op(e3, e3)) & (e1 = op(op(e3, op(e3, e3)), e3))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG131+1.p', ax14)).
fof(f890, plain, (~ spl15_106 | ~ spl15_21 | ~ spl15_99), inference(avatar_split_clause, [], [f294, f853, f397, f887])).
fof(f294, plain, (~ (e3 = op(e2, op(e2, e2))) | ~ (e0 = op(e2, e2)) | ~ (e1 = op(op(e2, op(e2, e2)), e2))), inference(cnf_transformation, [], [f37])).
fof(f37, plain, (~ (e3 = op(e2, op(e2, e2))) | ~ (e0 = op(e2, e2)) | ~ (e1 = op(op(e2, op(e2, e2)), e2))), inference(ennf_transformation, [], [f13])).
fof(f13, plain, ~ ((e3 = op(e2, op(e2, e2))) & (e0 = op(e2, e2)) & (e1 = op(op(e2, op(e2, e2)), e2))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG131+1.p', ax13)).
fof(f885, plain, (~ spl15_105 | ~ spl15_1 | ~ spl15_97), inference(avatar_split_clause, [], [f293, f844, f312, f882])).
fof(f293, plain, (~ (e2 = op(e3, op(e3, e3))) | ~ (e0 = op(e3, e3)) | ~ (e1 = op(op(e3, op(e3, e3)), e3))), inference(cnf_transformation, [], [f36])).
fof(f36, plain, (~ (e2 = op(e3, op(e3, e3))) | ~ (e0 = op(e3, e3)) | ~ (e1 = op(op(e3, op(e3, e3)), e3))), inference(ennf_transformation, [], [f12])).
fof(f12, plain, ~ ((e2 = op(e3, op(e3, e3))) & (e0 = op(e3, e3)) & (e1 = op(op(e3, op(e3, e3)), e3))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG131+1.p', ax12)).
fof(f880, plain, (~ spl15_101 | ~ spl15_44 | ~ spl15_104), inference(avatar_split_clause, [], [f292, f877, f494, f863])).
fof(f292, plain, (~ (e2 = op(e1, op(e1, e1))) | ~ (e3 = op(e1, e1)) | ~ (e0 = op(op(e1, op(e1, e1)), e1))), inference(cnf_transformation, [], [f35])).
fof(f35, plain, (~ (e2 = op(e1, op(e1, e1))) | ~ (e3 = op(e1, e1)) | ~ (e0 = op(op(e1, op(e1, e1)), e1))), inference(ennf_transformation, [], [f11])).
fof(f11, plain, ~ ((e2 = op(e1, op(e1, e1))) & (e3 = op(e1, e1)) & (e0 = op(op(e1, op(e1, e1)), e1))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG131+1.p', ax11)).
fof(f875, plain, (~ spl15_98 | ~ spl15_24 | ~ spl15_103), inference(avatar_split_clause, [], [f291, f872, f409, f849])).
fof(f291, plain, (~ (e1 = op(e2, op(e2, e2))) | ~ (e3 = op(e2, e2)) | ~ (e0 = op(op(e2, op(e2, e2)), e2))), inference(cnf_transformation, [], [f34])).
fof(f34, plain, (~ (e1 = op(e2, op(e2, e2))) | ~ (e3 = op(e2, e2)) | ~ (e0 = op(op(e2, op(e2, e2)), e2))), inference(ennf_transformation, [], [f10])).
fof(f10, plain, ~ ((e1 = op(e2, op(e2, e2))) & (e3 = op(e2, e2)) & (e0 = op(op(e2, op(e2, e2)), e2))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG131+1.p', ax10)).
fof(f870, plain, (~ spl15_101 | ~ spl15_43 | ~ spl15_102), inference(avatar_split_clause, [], [f290, f867, f490, f863])).
fof(f290, plain, (~ (e3 = op(e1, op(e1, e1))) | ~ (e2 = op(e1, e1)) | ~ (e0 = op(op(e1, op(e1, e1)), e1))), inference(cnf_transformation, [], [f33])).
fof(f33, plain, (~ (e3 = op(e1, op(e1, e1))) | ~ (e2 = op(e1, e1)) | ~ (e0 = op(op(e1, op(e1, e1)), e1))), inference(ennf_transformation, [], [f9])).
fof(f9, plain, ~ ((e3 = op(e1, op(e1, e1))) & (e2 = op(e1, e1)) & (e0 = op(op(e1, op(e1, e1)), e1))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG131+1.p', ax9)).
fof(f861, plain, (~ spl15_96 | ~ spl15_3 | ~ spl15_100), inference(avatar_split_clause, [], [f289, f858, f320, f840])).
fof(f289, plain, (~ (e1 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, e3)) | ~ (e0 = op(op(e3, op(e3, e3)), e3))), inference(cnf_transformation, [], [f32])).
fof(f32, plain, (~ (e1 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, e3)) | ~ (e0 = op(op(e3, op(e3, e3)), e3))), inference(ennf_transformation, [], [f8])).
fof(f8, plain, ~ ((e1 = op(e3, op(e3, e3))) & (e2 = op(e3, e3)) & (e0 = op(op(e3, op(e3, e3)), e3))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG131+1.p', ax8)).
fof(f856, plain, (~ spl15_98 | ~ spl15_22 | ~ spl15_99), inference(avatar_split_clause, [], [f288, f853, f401, f849])).
fof(f288, plain, (~ (e3 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, e2)) | ~ (e0 = op(op(e2, op(e2, e2)), e2))), inference(cnf_transformation, [], [f31])).
fof(f31, plain, (~ (e3 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, e2)) | ~ (e0 = op(op(e2, op(e2, e2)), e2))), inference(ennf_transformation, [], [f7])).
fof(f7, plain, ~ ((e3 = op(e2, op(e2, e2))) & (e1 = op(e2, e2)) & (e0 = op(op(e2, op(e2, e2)), e2))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG131+1.p', ax7)).
fof(f847, plain, (~ spl15_96 | ~ spl15_2 | ~ spl15_97), inference(avatar_split_clause, [], [f287, f844, f316, f840])).
fof(f287, plain, (~ (e2 = op(e3, op(e3, e3))) | ~ (e1 = op(e3, e3)) | ~ (e0 = op(op(e3, op(e3, e3)), e3))), inference(cnf_transformation, [], [f30])).
fof(f30, plain, (~ (e2 = op(e3, op(e3, e3))) | ~ (e1 = op(e3, e3)) | ~ (e0 = op(op(e3, op(e3, e3)), e3))), inference(ennf_transformation, [], [f6])).
fof(f6, plain, ~ ((e2 = op(e3, op(e3, e3))) & (e1 = op(e3, e3)) & (e0 = op(op(e3, op(e3, e3)), e3))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG131+1.p', ax6)).
fof(f838, plain, (spl15_92 | spl15_93 | spl15_94 | spl15_95), inference(avatar_split_clause, [], [f277, f835, f831, f827, f823])).
fof(f277, plain, ((e3 = op(e3, op(e3, e0))) | (e2 = op(e2, op(e2, e0))) | (e1 = op(e1, op(e1, e0))) | (e0 = op(e0, op(e0, e0)))), inference(cnf_transformation, [], [f69])).
fof(f69, plain, (((~ (e3 = op(e3, e3)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e1, e1)) & ~ (op(e0, e0) = e3) & (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0) & ((e3 = op(e3, op(e3, e3))) | (e2 = op(e2, op(e2, e3))) | (e1 = op(e1, op(e1, e3))) | (e0 = op(e0, op(e0, e3)))) & ((e3 = op(e3, op(e3, e2))) | (e2 = op(e2, op(e2, e2))) | (e1 = op(e1, op(e1, e2))) | (e0 = op(e0, op(e0, e2)))) & ((e3 = op(e3, op(e3, e1))) | (e2 = op(e2, op(e2, e1))) | (e1 = op(e1, op(e1, e1))) | (e0 = op(e0, op(e0, e1)))) & ((e3 = op(e3, op(e3, e0))) | (e2 = op(e2, op(e2, e0))) | (e1 = op(e1, op(e1, e0))) | (e0 = op(e0, op(e0, e0))))), inference(definition_folding, [], [f5, e68, e67, e66, e65, e64, e63, e62, e61, e60, e59, e58, e57, e56, e55, e54])).
fof(f54, plain, ((~ (e0 = op(e3, e3)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP0), inference(usedef, [], [e54])).
fof(e54, plain, (sP0 <=> (~ (e0 = op(e3, e3)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f55, plain, ((~ (e1 = op(e3, e3)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e1, e1)) & ~ (op(e0, e0) = e1) & (e0 = op(e1, e0)) & (e0 = op(e0, e1))) | ~ sP1), inference(usedef, [], [e55])).
fof(e55, plain, (sP1 <=> (~ (e1 = op(e3, e3)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e1, e1)) & ~ (op(e0, e0) = e1) & (e0 = op(e1, e0)) & (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f56, plain, ((~ (e2 = op(e3, e3)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e1, e1)) & ~ (op(e0, e0) = e2) & (e0 = op(e2, e0)) & (e0 = op(e0, e2))) | ~ sP2), inference(usedef, [], [e56])).
fof(e56, plain, (sP2 <=> (~ (e2 = op(e3, e3)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e1, e1)) & ~ (op(e0, e0) = e2) & (e0 = op(e2, e0)) & (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f57, plain, ((~ (e3 = op(e3, e3)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e1, e1)) & ~ (op(e0, e0) = e3) & (e0 = op(e3, e0)) & (e0 = op(e0, e3))) | ~ sP3), inference(usedef, [], [e57])).
fof(e57, plain, (sP3 <=> (~ (e3 = op(e3, e3)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e1, e1)) & ~ (op(e0, e0) = e3) & (e0 = op(e3, e0)) & (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f58, plain, ((~ (e0 = op(e3, e3)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e0, e0)) & (e1 = op(e0, e1)) & (e1 = op(e1, e0))) | ~ sP4), inference(usedef, [], [e58])).
fof(e58, plain, (sP4 <=> (~ (e0 = op(e3, e3)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e0, e0)) & (e1 = op(e0, e1)) & (e1 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f59, plain, ((~ (e1 = op(e3, e3)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e1, e1)) & ~ (op(e0, e0) = e1) & (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP5), inference(usedef, [], [e59])).
fof(e59, plain, (sP5 <=> (~ (e1 = op(e3, e3)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e1, e1)) & ~ (op(e0, e0) = e1) & (e1 = op(e1, e1)) & (e1 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f60, plain, ((~ (e2 = op(e3, e3)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e1, e1)) & ~ (op(e0, e0) = e2) & (e1 = op(e2, e1)) & (e1 = op(e1, e2))) | ~ sP6), inference(usedef, [], [e60])).
fof(e60, plain, (sP6 <=> (~ (e2 = op(e3, e3)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e1, e1)) & ~ (op(e0, e0) = e2) & (e1 = op(e2, e1)) & (e1 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f61, plain, ((~ (e3 = op(e3, e3)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e1, e1)) & ~ (op(e0, e0) = e3) & (e1 = op(e3, e1)) & (e1 = op(e1, e3))) | ~ sP7), inference(usedef, [], [e61])).
fof(e61, plain, (sP7 <=> (~ (e3 = op(e3, e3)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e1, e1)) & ~ (op(e0, e0) = e3) & (e1 = op(e3, e1)) & (e1 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f62, plain, ((~ (e0 = op(e3, e3)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e0, e0)) & (e2 = op(e0, e2)) & (e2 = op(e2, e0))) | ~ sP8), inference(usedef, [], [e62])).
fof(e62, plain, (sP8 <=> (~ (e0 = op(e3, e3)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e0, e0)) & (e2 = op(e0, e2)) & (e2 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f63, plain, ((~ (e1 = op(e3, e3)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e1, e1)) & ~ (op(e0, e0) = e1) & (e2 = op(e1, e2)) & (e2 = op(e2, e1))) | ~ sP9), inference(usedef, [], [e63])).
fof(e63, plain, (sP9 <=> (~ (e1 = op(e3, e3)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e1, e1)) & ~ (op(e0, e0) = e1) & (e2 = op(e1, e2)) & (e2 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f64, plain, ((~ (e2 = op(e3, e3)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e1, e1)) & ~ (op(e0, e0) = e2) & (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP10), inference(usedef, [], [e64])).
fof(e64, plain, (sP10 <=> (~ (e2 = op(e3, e3)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e1, e1)) & ~ (op(e0, e0) = e2) & (e2 = op(e2, e2)) & (e2 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f65, plain, ((~ (e3 = op(e3, e3)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e1, e1)) & ~ (op(e0, e0) = e3) & (e2 = op(e3, e2)) & (e2 = op(e2, e3))) | ~ sP11), inference(usedef, [], [e65])).
fof(e65, plain, (sP11 <=> (~ (e3 = op(e3, e3)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e1, e1)) & ~ (op(e0, e0) = e3) & (e2 = op(e3, e2)) & (e2 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f66, plain, ((~ (e0 = op(e3, e3)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e0, e0)) & (e3 = op(e0, e3)) & (e3 = op(e3, e0))) | ~ sP12), inference(usedef, [], [e66])).
fof(e66, plain, (sP12 <=> (~ (e0 = op(e3, e3)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e0, e0)) & (e3 = op(e0, e3)) & (e3 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f67, plain, ((~ (e1 = op(e3, e3)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e1, e1)) & ~ (op(e0, e0) = e1) & (e3 = op(e1, e3)) & (e3 = op(e3, e1))) | ~ sP13), inference(usedef, [], [e67])).
fof(e67, plain, (sP13 <=> (~ (e1 = op(e3, e3)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e1, e1)) & ~ (op(e0, e0) = e1) & (e3 = op(e1, e3)) & (e3 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f68, plain, ((~ (e2 = op(e3, e3)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e1, e1)) & ~ (op(e0, e0) = e2) & (e3 = op(e2, e3)) & (e3 = op(e3, e2))) | ~ sP14), inference(usedef, [], [e68])).
fof(e68, plain, (sP14 <=> (~ (e2 = op(e3, e3)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e1, e1)) & ~ (op(e0, e0) = e2) & (e3 = op(e2, e3)) & (e3 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f5, plain, (((~ (e3 = op(e3, e3)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e1, e1)) & ~ (op(e0, e0) = e3) & (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | (~ (e2 = op(e3, e3)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e1, e1)) & ~ (op(e0, e0) = e2) & (e3 = op(e2, e3)) & (e3 = op(e3, e2))) | (~ (e1 = op(e3, e3)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e1, e1)) & ~ (op(e0, e0) = e1) & (e3 = op(e1, e3)) & (e3 = op(e3, e1))) | (~ (e0 = op(e3, e3)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e0, e0)) & (e3 = op(e0, e3)) & (e3 = op(e3, e0))) | (~ (e3 = op(e3, e3)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e1, e1)) & ~ (op(e0, e0) = e3) & (e2 = op(e3, e2)) & (e2 = op(e2, e3))) | (~ (e2 = op(e3, e3)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e1, e1)) & ~ (op(e0, e0) = e2) & (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | (~ (e1 = op(e3, e3)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e1, e1)) & ~ (op(e0, e0) = e1) & (e2 = op(e1, e2)) & (e2 = op(e2, e1))) | (~ (e0 = op(e3, e3)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e0, e0)) & (e2 = op(e0, e2)) & (e2 = op(e2, e0))) | (~ (e3 = op(e3, e3)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e1, e1)) & ~ (op(e0, e0) = e3) & (e1 = op(e3, e1)) & (e1 = op(e1, e3))) | (~ (e2 = op(e3, e3)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e1, e1)) & ~ (op(e0, e0) = e2) & (e1 = op(e2, e1)) & (e1 = op(e1, e2))) | (~ (e1 = op(e3, e3)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e1, e1)) & ~ (op(e0, e0) = e1) & (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | (~ (e0 = op(e3, e3)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e0, e0)) & (e1 = op(e0, e1)) & (e1 = op(e1, e0))) | (~ (e3 = op(e3, e3)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e1, e1)) & ~ (op(e0, e0) = e3) & (e0 = op(e3, e0)) & (e0 = op(e0, e3))) | (~ (e2 = op(e3, e3)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e1, e1)) & ~ (op(e0, e0) = e2) & (e0 = op(e2, e0)) & (e0 = op(e0, e2))) | (~ (e1 = op(e3, e3)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e1, e1)) & ~ (op(e0, e0) = e1) & (e0 = op(e1, e0)) & (e0 = op(e0, e1))) | (~ (e0 = op(e3, e3)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0)))) & ((e3 = op(e3, op(e3, e3))) | (e2 = op(e2, op(e2, e3))) | (e1 = op(e1, op(e1, e3))) | (e0 = op(e0, op(e0, e3)))) & ((e3 = op(e3, op(e3, e2))) | (e2 = op(e2, op(e2, e2))) | (e1 = op(e1, op(e1, e2))) | (e0 = op(e0, op(e0, e2)))) & ((e3 = op(e3, op(e3, e1))) | (e2 = op(e2, op(e2, e1))) | (e1 = op(e1, op(e1, e1))) | (e0 = op(e0, op(e0, e1)))) & ((e3 = op(e3, op(e3, e0))) | (e2 = op(e2, op(e2, e0))) | (e1 = op(e1, op(e1, e0))) | (e0 = op(e0, op(e0, e0))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG131+1.p', ax5)).
fof(f821, plain, (spl15_88 | spl15_89 | spl15_90 | spl15_91), inference(avatar_split_clause, [], [f278, f818, f814, f810, f806])).
fof(f278, plain, ((e3 = op(e3, op(e3, e1))) | (e2 = op(e2, op(e2, e1))) | (e1 = op(e1, op(e1, e1))) | (e0 = op(e0, op(e0, e1)))), inference(cnf_transformation, [], [f69])).
fof(f804, plain, (spl15_84 | spl15_85 | spl15_86 | spl15_87), inference(avatar_split_clause, [], [f279, f801, f797, f793, f789])).
fof(f279, plain, ((e3 = op(e3, op(e3, e2))) | (e2 = op(e2, op(e2, e2))) | (e1 = op(e1, op(e1, e2))) | (e0 = op(e0, op(e0, e2)))), inference(cnf_transformation, [], [f69])).
fof(f787, plain, (spl15_80 | spl15_81 | spl15_82 | spl15_83), inference(avatar_split_clause, [], [f280, f784, f780, f776, f772])).
fof(f280, plain, ((e3 = op(e3, op(e3, e3))) | (e2 = op(e2, op(e2, e3))) | (e1 = op(e1, op(e1, e3))) | (e0 = op(e0, op(e0, e3)))), inference(cnf_transformation, [], [f69])).
fof(f770, plain, (spl15_79 | spl15_78 | spl15_77 | spl15_76 | spl15_75 | spl15_74 | spl15_73 | spl15_72 | spl15_71 | spl15_70 | spl15_69 | spl15_68 | spl15_67 | spl15_66 | spl15_65 | spl15_4), inference(avatar_split_clause, [], [f281, f324, f616, f626, f636, f646, f656, f666, f676, f686, f696, f706, f716, f726, f736, f746, f756])).
fof(f756, plain, (spl15_79 <=> sP0), introduced(avatar_definition, [new_symbols(naming, [spl15_79])])).
fof(f746, plain, (spl15_78 <=> sP1), introduced(avatar_definition, [new_symbols(naming, [spl15_78])])).
fof(f736, plain, (spl15_77 <=> sP2), introduced(avatar_definition, [new_symbols(naming, [spl15_77])])).
fof(f726, plain, (spl15_76 <=> sP3), introduced(avatar_definition, [new_symbols(naming, [spl15_76])])).
fof(f716, plain, (spl15_75 <=> sP4), introduced(avatar_definition, [new_symbols(naming, [spl15_75])])).
fof(f706, plain, (spl15_74 <=> sP5), introduced(avatar_definition, [new_symbols(naming, [spl15_74])])).
fof(f696, plain, (spl15_73 <=> sP6), introduced(avatar_definition, [new_symbols(naming, [spl15_73])])).
fof(f686, plain, (spl15_72 <=> sP7), introduced(avatar_definition, [new_symbols(naming, [spl15_72])])).
fof(f676, plain, (spl15_71 <=> sP8), introduced(avatar_definition, [new_symbols(naming, [spl15_71])])).
fof(f666, plain, (spl15_70 <=> sP9), introduced(avatar_definition, [new_symbols(naming, [spl15_70])])).
fof(f656, plain, (spl15_69 <=> sP10), introduced(avatar_definition, [new_symbols(naming, [spl15_69])])).
fof(f646, plain, (spl15_68 <=> sP11), introduced(avatar_definition, [new_symbols(naming, [spl15_68])])).
fof(f636, plain, (spl15_67 <=> sP12), introduced(avatar_definition, [new_symbols(naming, [spl15_67])])).
fof(f626, plain, (spl15_66 <=> sP13), introduced(avatar_definition, [new_symbols(naming, [spl15_66])])).
fof(f616, plain, (spl15_65 <=> sP14), introduced(avatar_definition, [new_symbols(naming, [spl15_65])])).
fof(f281, plain, ((e3 = op(e3, e3)) | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f69])).
fof(f766, plain, (spl15_79 | spl15_78 | spl15_77 | spl15_76 | spl15_75 | spl15_74 | spl15_73 | spl15_72 | spl15_71 | spl15_70 | spl15_69 | spl15_68 | spl15_67 | spl15_66 | spl15_65 | ~ spl15_24), inference(avatar_split_clause, [], [f285, f409, f616, f626, f636, f646, f656, f666, f676, f686, f696, f706, f716, f726, f736, f746, f756])).
fof(f285, plain, (~ (e3 = op(e2, e2)) | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f69])).
fof(f765, plain, (spl15_79 | spl15_78 | spl15_77 | spl15_76 | spl15_75 | spl15_74 | spl15_73 | spl15_72 | spl15_71 | spl15_70 | spl15_69 | spl15_68 | spl15_67 | spl15_66 | spl15_65 | ~ spl15_4), inference(avatar_split_clause, [], [f286, f324, f616, f626, f636, f646, f656, f666, f676, f686, f696, f706, f716, f726, f736, f746, f756])).
fof(f286, plain, (~ (e3 = op(e3, e3)) | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f69])).
fof(f764, plain, (~ spl15_79 | spl15_61), inference(avatar_split_clause, [], [f271, f567, f756])).
fof(f271, plain, ((e0 = op(e0, e0)) | ~ sP0), inference(cnf_transformation, [], [f84])).
fof(f84, plain, ((~ (e0 = op(e3, e3)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP0), inference(nnf_transformation, [], [f54])).
fof(f762, plain, (~ spl15_79 | ~ spl15_61), inference(avatar_split_clause, [], [f273, f567, f756])).
fof(f273, plain, (~ (e0 = op(e0, e0)) | ~ sP0), inference(cnf_transformation, [], [f84])).
fof(f754, plain, (~ spl15_78 | spl15_57), inference(avatar_split_clause, [], [f265, f550, f746])).
fof(f265, plain, ((e0 = op(e0, e1)) | ~ sP1), inference(cnf_transformation, [], [f83])).
fof(f83, plain, ((~ (e1 = op(e3, e3)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e1, e1)) & ~ (op(e0, e0) = e1) & (e0 = op(e1, e0)) & (e0 = op(e0, e1))) | ~ sP1), inference(nnf_transformation, [], [f55])).
fof(f753, plain, (~ spl15_78 | spl15_45), inference(avatar_split_clause, [], [f266, f499, f746])).
fof(f266, plain, ((e0 = op(e1, e0)) | ~ sP1), inference(cnf_transformation, [], [f83])).
fof(f752, plain, (~ spl15_78 | ~ spl15_62), inference(avatar_split_clause, [], [f267, f571, f746])).
fof(f267, plain, (~ (op(e0, e0) = e1) | ~ sP1), inference(cnf_transformation, [], [f83])).
fof(f751, plain, (~ spl15_78 | ~ spl15_42), inference(avatar_split_clause, [], [f268, f486, f746])).
fof(f268, plain, (~ (e1 = op(e1, e1)) | ~ sP1), inference(cnf_transformation, [], [f83])).
fof(f750, plain, (~ spl15_78 | ~ spl15_22), inference(avatar_split_clause, [], [f269, f401, f746])).
fof(f269, plain, (~ (e1 = op(e2, e2)) | ~ sP1), inference(cnf_transformation, [], [f83])).
fof(f749, plain, (~ spl15_78 | ~ spl15_2), inference(avatar_split_clause, [], [f270, f316, f746])).
fof(f270, plain, (~ (e1 = op(e3, e3)) | ~ sP1), inference(cnf_transformation, [], [f83])).
fof(f744, plain, (~ spl15_77 | spl15_53), inference(avatar_split_clause, [], [f259, f533, f736])).
fof(f259, plain, ((e0 = op(e0, e2)) | ~ sP2), inference(cnf_transformation, [], [f82])).
fof(f82, plain, ((~ (e2 = op(e3, e3)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e1, e1)) & ~ (op(e0, e0) = e2) & (e0 = op(e2, e0)) & (e0 = op(e0, e2))) | ~ sP2), inference(nnf_transformation, [], [f56])).
fof(f743, plain, (~ spl15_77 | spl15_29), inference(avatar_split_clause, [], [f260, f431, f736])).
fof(f260, plain, ((e0 = op(e2, e0)) | ~ sP2), inference(cnf_transformation, [], [f82])).
fof(f742, plain, (~ spl15_77 | ~ spl15_63), inference(avatar_split_clause, [], [f261, f575, f736])).
fof(f261, plain, (~ (op(e0, e0) = e2) | ~ sP2), inference(cnf_transformation, [], [f82])).
fof(f741, plain, (~ spl15_77 | ~ spl15_43), inference(avatar_split_clause, [], [f262, f490, f736])).
fof(f262, plain, (~ (e2 = op(e1, e1)) | ~ sP2), inference(cnf_transformation, [], [f82])).
fof(f740, plain, (~ spl15_77 | ~ spl15_23), inference(avatar_split_clause, [], [f263, f405, f736])).
fof(f263, plain, (~ (e2 = op(e2, e2)) | ~ sP2), inference(cnf_transformation, [], [f82])).
fof(f734, plain, (~ spl15_76 | spl15_49), inference(avatar_split_clause, [], [f253, f516, f726])).
fof(f253, plain, ((e0 = op(e0, e3)) | ~ sP3), inference(cnf_transformation, [], [f81])).
fof(f81, plain, ((~ (e3 = op(e3, e3)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e1, e1)) & ~ (op(e0, e0) = e3) & (e0 = op(e3, e0)) & (e0 = op(e0, e3))) | ~ sP3), inference(nnf_transformation, [], [f57])).
fof(f733, plain, (~ spl15_76 | spl15_13), inference(avatar_split_clause, [], [f254, f363, f726])).
fof(f254, plain, ((e0 = op(e3, e0)) | ~ sP3), inference(cnf_transformation, [], [f81])).
fof(f732, plain, (~ spl15_76 | ~ spl15_64), inference(avatar_split_clause, [], [f255, f579, f726])).
fof(f255, plain, (~ (op(e0, e0) = e3) | ~ sP3), inference(cnf_transformation, [], [f81])).
fof(f731, plain, (~ spl15_76 | ~ spl15_44), inference(avatar_split_clause, [], [f256, f494, f726])).
fof(f256, plain, (~ (e3 = op(e1, e1)) | ~ sP3), inference(cnf_transformation, [], [f81])).
fof(f730, plain, (~ spl15_76 | ~ spl15_24), inference(avatar_split_clause, [], [f257, f409, f726])).
fof(f257, plain, (~ (e3 = op(e2, e2)) | ~ sP3), inference(cnf_transformation, [], [f81])).
fof(f729, plain, (~ spl15_76 | ~ spl15_4), inference(avatar_split_clause, [], [f258, f324, f726])).
fof(f258, plain, (~ (e3 = op(e3, e3)) | ~ sP3), inference(cnf_transformation, [], [f81])).
fof(f724, plain, (~ spl15_75 | spl15_46), inference(avatar_split_clause, [], [f247, f503, f716])).
fof(f247, plain, ((e1 = op(e1, e0)) | ~ sP4), inference(cnf_transformation, [], [f80])).
fof(f80, plain, ((~ (e0 = op(e3, e3)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e0, e0)) & (e1 = op(e0, e1)) & (e1 = op(e1, e0))) | ~ sP4), inference(nnf_transformation, [], [f58])).
fof(f723, plain, (~ spl15_75 | spl15_58), inference(avatar_split_clause, [], [f248, f554, f716])).
fof(f248, plain, ((e1 = op(e0, e1)) | ~ sP4), inference(cnf_transformation, [], [f80])).
fof(f722, plain, (~ spl15_75 | ~ spl15_61), inference(avatar_split_clause, [], [f249, f567, f716])).
fof(f249, plain, (~ (e0 = op(e0, e0)) | ~ sP4), inference(cnf_transformation, [], [f80])).
fof(f721, plain, (~ spl15_75 | ~ spl15_41), inference(avatar_split_clause, [], [f250, f482, f716])).
fof(f250, plain, (~ (e0 = op(e1, e1)) | ~ sP4), inference(cnf_transformation, [], [f80])).
fof(f720, plain, (~ spl15_75 | ~ spl15_21), inference(avatar_split_clause, [], [f251, f397, f716])).
fof(f251, plain, (~ (e0 = op(e2, e2)) | ~ sP4), inference(cnf_transformation, [], [f80])).
fof(f719, plain, (~ spl15_75 | ~ spl15_1), inference(avatar_split_clause, [], [f252, f312, f716])).
fof(f252, plain, (~ (e0 = op(e3, e3)) | ~ sP4), inference(cnf_transformation, [], [f80])).
fof(f714, plain, (~ spl15_74 | spl15_42), inference(avatar_split_clause, [], [f241, f486, f706])).
fof(f241, plain, ((e1 = op(e1, e1)) | ~ sP5), inference(cnf_transformation, [], [f79])).
fof(f79, plain, ((~ (e1 = op(e3, e3)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e1, e1)) & ~ (op(e0, e0) = e1) & (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP5), inference(nnf_transformation, [], [f59])).
fof(f711, plain, (~ spl15_74 | ~ spl15_42), inference(avatar_split_clause, [], [f244, f486, f706])).
fof(f244, plain, (~ (e1 = op(e1, e1)) | ~ sP5), inference(cnf_transformation, [], [f79])).
fof(f704, plain, (~ spl15_73 | spl15_38), inference(avatar_split_clause, [], [f235, f469, f696])).
fof(f235, plain, ((e1 = op(e1, e2)) | ~ sP6), inference(cnf_transformation, [], [f78])).
fof(f78, plain, ((~ (e2 = op(e3, e3)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e1, e1)) & ~ (op(e0, e0) = e2) & (e1 = op(e2, e1)) & (e1 = op(e1, e2))) | ~ sP6), inference(nnf_transformation, [], [f60])).
fof(f703, plain, (~ spl15_73 | spl15_26), inference(avatar_split_clause, [], [f236, f418, f696])).
fof(f236, plain, ((e1 = op(e2, e1)) | ~ sP6), inference(cnf_transformation, [], [f78])).
fof(f702, plain, (~ spl15_73 | ~ spl15_63), inference(avatar_split_clause, [], [f237, f575, f696])).
fof(f237, plain, (~ (op(e0, e0) = e2) | ~ sP6), inference(cnf_transformation, [], [f78])).
fof(f701, plain, (~ spl15_73 | ~ spl15_43), inference(avatar_split_clause, [], [f238, f490, f696])).
fof(f238, plain, (~ (e2 = op(e1, e1)) | ~ sP6), inference(cnf_transformation, [], [f78])).
fof(f700, plain, (~ spl15_73 | ~ spl15_23), inference(avatar_split_clause, [], [f239, f405, f696])).
fof(f239, plain, (~ (e2 = op(e2, e2)) | ~ sP6), inference(cnf_transformation, [], [f78])).
fof(f699, plain, (~ spl15_73 | ~ spl15_3), inference(avatar_split_clause, [], [f240, f320, f696])).
fof(f240, plain, (~ (e2 = op(e3, e3)) | ~ sP6), inference(cnf_transformation, [], [f78])).
fof(f694, plain, (~ spl15_72 | spl15_34), inference(avatar_split_clause, [], [f229, f452, f686])).
fof(f229, plain, ((e1 = op(e1, e3)) | ~ sP7), inference(cnf_transformation, [], [f77])).
fof(f77, plain, ((~ (e3 = op(e3, e3)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e1, e1)) & ~ (op(e0, e0) = e3) & (e1 = op(e3, e1)) & (e1 = op(e1, e3))) | ~ sP7), inference(nnf_transformation, [], [f61])).
fof(f693, plain, (~ spl15_72 | spl15_10), inference(avatar_split_clause, [], [f230, f350, f686])).
fof(f230, plain, ((e1 = op(e3, e1)) | ~ sP7), inference(cnf_transformation, [], [f77])).
fof(f692, plain, (~ spl15_72 | ~ spl15_64), inference(avatar_split_clause, [], [f231, f579, f686])).
fof(f231, plain, (~ (op(e0, e0) = e3) | ~ sP7), inference(cnf_transformation, [], [f77])).
fof(f691, plain, (~ spl15_72 | ~ spl15_44), inference(avatar_split_clause, [], [f232, f494, f686])).
fof(f232, plain, (~ (e3 = op(e1, e1)) | ~ sP7), inference(cnf_transformation, [], [f77])).
fof(f690, plain, (~ spl15_72 | ~ spl15_24), inference(avatar_split_clause, [], [f233, f409, f686])).
fof(f233, plain, (~ (e3 = op(e2, e2)) | ~ sP7), inference(cnf_transformation, [], [f77])).
fof(f689, plain, (~ spl15_72 | ~ spl15_4), inference(avatar_split_clause, [], [f234, f324, f686])).
fof(f234, plain, (~ (e3 = op(e3, e3)) | ~ sP7), inference(cnf_transformation, [], [f77])).
fof(f684, plain, (~ spl15_71 | spl15_31), inference(avatar_split_clause, [], [f223, f439, f676])).
fof(f223, plain, ((e2 = op(e2, e0)) | ~ sP8), inference(cnf_transformation, [], [f76])).
fof(f76, plain, ((~ (e0 = op(e3, e3)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e0, e0)) & (e2 = op(e0, e2)) & (e2 = op(e2, e0))) | ~ sP8), inference(nnf_transformation, [], [f62])).
fof(f683, plain, (~ spl15_71 | spl15_55), inference(avatar_split_clause, [], [f224, f541, f676])).
fof(f224, plain, ((e2 = op(e0, e2)) | ~ sP8), inference(cnf_transformation, [], [f76])).
fof(f682, plain, (~ spl15_71 | ~ spl15_61), inference(avatar_split_clause, [], [f225, f567, f676])).
fof(f225, plain, (~ (e0 = op(e0, e0)) | ~ sP8), inference(cnf_transformation, [], [f76])).
fof(f681, plain, (~ spl15_71 | ~ spl15_41), inference(avatar_split_clause, [], [f226, f482, f676])).
fof(f226, plain, (~ (e0 = op(e1, e1)) | ~ sP8), inference(cnf_transformation, [], [f76])).
fof(f680, plain, (~ spl15_71 | ~ spl15_21), inference(avatar_split_clause, [], [f227, f397, f676])).
fof(f227, plain, (~ (e0 = op(e2, e2)) | ~ sP8), inference(cnf_transformation, [], [f76])).
fof(f679, plain, (~ spl15_71 | ~ spl15_1), inference(avatar_split_clause, [], [f228, f312, f676])).
fof(f228, plain, (~ (e0 = op(e3, e3)) | ~ sP8), inference(cnf_transformation, [], [f76])).
fof(f674, plain, (~ spl15_70 | spl15_27), inference(avatar_split_clause, [], [f217, f422, f666])).
fof(f217, plain, ((e2 = op(e2, e1)) | ~ sP9), inference(cnf_transformation, [], [f75])).
fof(f75, plain, ((~ (e1 = op(e3, e3)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e1, e1)) & ~ (op(e0, e0) = e1) & (e2 = op(e1, e2)) & (e2 = op(e2, e1))) | ~ sP9), inference(nnf_transformation, [], [f63])).
fof(f673, plain, (~ spl15_70 | spl15_39), inference(avatar_split_clause, [], [f218, f473, f666])).
fof(f218, plain, ((e2 = op(e1, e2)) | ~ sP9), inference(cnf_transformation, [], [f75])).
fof(f672, plain, (~ spl15_70 | ~ spl15_62), inference(avatar_split_clause, [], [f219, f571, f666])).
fof(f219, plain, (~ (op(e0, e0) = e1) | ~ sP9), inference(cnf_transformation, [], [f75])).
fof(f671, plain, (~ spl15_70 | ~ spl15_42), inference(avatar_split_clause, [], [f220, f486, f666])).
fof(f220, plain, (~ (e1 = op(e1, e1)) | ~ sP9), inference(cnf_transformation, [], [f75])).
fof(f670, plain, (~ spl15_70 | ~ spl15_22), inference(avatar_split_clause, [], [f221, f401, f666])).
fof(f221, plain, (~ (e1 = op(e2, e2)) | ~ sP9), inference(cnf_transformation, [], [f75])).
fof(f669, plain, (~ spl15_70 | ~ spl15_2), inference(avatar_split_clause, [], [f222, f316, f666])).
fof(f222, plain, (~ (e1 = op(e3, e3)) | ~ sP9), inference(cnf_transformation, [], [f75])).
fof(f664, plain, (~ spl15_69 | spl15_23), inference(avatar_split_clause, [], [f211, f405, f656])).
fof(f211, plain, ((e2 = op(e2, e2)) | ~ sP10), inference(cnf_transformation, [], [f74])).
fof(f74, plain, ((~ (e2 = op(e3, e3)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e1, e1)) & ~ (op(e0, e0) = e2) & (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP10), inference(nnf_transformation, [], [f64])).
fof(f660, plain, (~ spl15_69 | ~ spl15_23), inference(avatar_split_clause, [], [f215, f405, f656])).
fof(f215, plain, (~ (e2 = op(e2, e2)) | ~ sP10), inference(cnf_transformation, [], [f74])).
fof(f654, plain, (~ spl15_68 | spl15_19), inference(avatar_split_clause, [], [f205, f388, f646])).
fof(f205, plain, ((e2 = op(e2, e3)) | ~ sP11), inference(cnf_transformation, [], [f73])).
fof(f73, plain, ((~ (e3 = op(e3, e3)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e1, e1)) & ~ (op(e0, e0) = e3) & (e2 = op(e3, e2)) & (e2 = op(e2, e3))) | ~ sP11), inference(nnf_transformation, [], [f65])).
fof(f653, plain, (~ spl15_68 | spl15_7), inference(avatar_split_clause, [], [f206, f337, f646])).
fof(f206, plain, ((e2 = op(e3, e2)) | ~ sP11), inference(cnf_transformation, [], [f73])).
fof(f652, plain, (~ spl15_68 | ~ spl15_64), inference(avatar_split_clause, [], [f207, f579, f646])).
fof(f207, plain, (~ (op(e0, e0) = e3) | ~ sP11), inference(cnf_transformation, [], [f73])).
fof(f651, plain, (~ spl15_68 | ~ spl15_44), inference(avatar_split_clause, [], [f208, f494, f646])).
fof(f208, plain, (~ (e3 = op(e1, e1)) | ~ sP11), inference(cnf_transformation, [], [f73])).
fof(f650, plain, (~ spl15_68 | ~ spl15_24), inference(avatar_split_clause, [], [f209, f409, f646])).
fof(f209, plain, (~ (e3 = op(e2, e2)) | ~ sP11), inference(cnf_transformation, [], [f73])).
fof(f649, plain, (~ spl15_68 | ~ spl15_4), inference(avatar_split_clause, [], [f210, f324, f646])).
fof(f210, plain, (~ (e3 = op(e3, e3)) | ~ sP11), inference(cnf_transformation, [], [f73])).
fof(f644, plain, (~ spl15_67 | spl15_16), inference(avatar_split_clause, [], [f199, f375, f636])).
fof(f199, plain, ((e3 = op(e3, e0)) | ~ sP12), inference(cnf_transformation, [], [f72])).
fof(f72, plain, ((~ (e0 = op(e3, e3)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e0, e0)) & (e3 = op(e0, e3)) & (e3 = op(e3, e0))) | ~ sP12), inference(nnf_transformation, [], [f66])).
fof(f643, plain, (~ spl15_67 | spl15_52), inference(avatar_split_clause, [], [f200, f528, f636])).
fof(f200, plain, ((e3 = op(e0, e3)) | ~ sP12), inference(cnf_transformation, [], [f72])).
fof(f642, plain, (~ spl15_67 | ~ spl15_61), inference(avatar_split_clause, [], [f201, f567, f636])).
fof(f201, plain, (~ (e0 = op(e0, e0)) | ~ sP12), inference(cnf_transformation, [], [f72])).
fof(f641, plain, (~ spl15_67 | ~ spl15_41), inference(avatar_split_clause, [], [f202, f482, f636])).
fof(f202, plain, (~ (e0 = op(e1, e1)) | ~ sP12), inference(cnf_transformation, [], [f72])).
fof(f640, plain, (~ spl15_67 | ~ spl15_21), inference(avatar_split_clause, [], [f203, f397, f636])).
fof(f203, plain, (~ (e0 = op(e2, e2)) | ~ sP12), inference(cnf_transformation, [], [f72])).
fof(f639, plain, (~ spl15_67 | ~ spl15_1), inference(avatar_split_clause, [], [f204, f312, f636])).
fof(f204, plain, (~ (e0 = op(e3, e3)) | ~ sP12), inference(cnf_transformation, [], [f72])).
fof(f634, plain, (~ spl15_66 | spl15_12), inference(avatar_split_clause, [], [f193, f358, f626])).
fof(f193, plain, ((e3 = op(e3, e1)) | ~ sP13), inference(cnf_transformation, [], [f71])).
fof(f71, plain, ((~ (e1 = op(e3, e3)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e1, e1)) & ~ (op(e0, e0) = e1) & (e3 = op(e1, e3)) & (e3 = op(e3, e1))) | ~ sP13), inference(nnf_transformation, [], [f67])).
fof(f633, plain, (~ spl15_66 | spl15_36), inference(avatar_split_clause, [], [f194, f460, f626])).
fof(f194, plain, ((e3 = op(e1, e3)) | ~ sP13), inference(cnf_transformation, [], [f71])).
fof(f632, plain, (~ spl15_66 | ~ spl15_62), inference(avatar_split_clause, [], [f195, f571, f626])).
fof(f195, plain, (~ (op(e0, e0) = e1) | ~ sP13), inference(cnf_transformation, [], [f71])).
fof(f631, plain, (~ spl15_66 | ~ spl15_42), inference(avatar_split_clause, [], [f196, f486, f626])).
fof(f196, plain, (~ (e1 = op(e1, e1)) | ~ sP13), inference(cnf_transformation, [], [f71])).
fof(f630, plain, (~ spl15_66 | ~ spl15_22), inference(avatar_split_clause, [], [f197, f401, f626])).
fof(f197, plain, (~ (e1 = op(e2, e2)) | ~ sP13), inference(cnf_transformation, [], [f71])).
fof(f629, plain, (~ spl15_66 | ~ spl15_2), inference(avatar_split_clause, [], [f198, f316, f626])).
fof(f198, plain, (~ (e1 = op(e3, e3)) | ~ sP13), inference(cnf_transformation, [], [f71])).
fof(f624, plain, (~ spl15_65 | spl15_8), inference(avatar_split_clause, [], [f187, f341, f616])).
fof(f187, plain, ((e3 = op(e3, e2)) | ~ sP14), inference(cnf_transformation, [], [f70])).
fof(f70, plain, ((~ (e2 = op(e3, e3)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e1, e1)) & ~ (op(e0, e0) = e2) & (e3 = op(e2, e3)) & (e3 = op(e3, e2))) | ~ sP14), inference(nnf_transformation, [], [f68])).
fof(f623, plain, (~ spl15_65 | spl15_20), inference(avatar_split_clause, [], [f188, f392, f616])).
fof(f188, plain, ((e3 = op(e2, e3)) | ~ sP14), inference(cnf_transformation, [], [f70])).
fof(f622, plain, (~ spl15_65 | ~ spl15_63), inference(avatar_split_clause, [], [f189, f575, f616])).
fof(f189, plain, (~ (op(e0, e0) = e2) | ~ sP14), inference(cnf_transformation, [], [f70])).
fof(f621, plain, (~ spl15_65 | ~ spl15_43), inference(avatar_split_clause, [], [f190, f490, f616])).
fof(f190, plain, (~ (e2 = op(e1, e1)) | ~ sP14), inference(cnf_transformation, [], [f70])).
fof(f620, plain, (~ spl15_65 | ~ spl15_23), inference(avatar_split_clause, [], [f191, f405, f616])).
fof(f191, plain, (~ (e2 = op(e2, e2)) | ~ sP14), inference(cnf_transformation, [], [f70])).
fof(f619, plain, (~ spl15_65 | ~ spl15_3), inference(avatar_split_clause, [], [f192, f320, f616])).
fof(f192, plain, (~ (e2 = op(e3, e3)) | ~ sP14), inference(cnf_transformation, [], [f70])).
fof(f614, plain, (spl15_61 | spl15_57 | spl15_53 | spl15_49), inference(avatar_split_clause, [], [f101, f516, f533, f550, f567])).
fof(f101, plain, ((e0 = op(e0, e3)) | (e0 = op(e0, e2)) | (e0 = op(e0, e1)) | (e0 = op(e0, e0))), inference(cnf_transformation, [], [f2])).
fof(f2, plain, (((e3 = op(e3, e3)) | (e3 = op(e2, e3)) | (e3 = op(e1, e3)) | (e3 = op(e0, e3))) & ((e3 = op(e3, e3)) | (e3 = op(e3, e2)) | (e3 = op(e3, e1)) | (e3 = op(e3, e0))) & ((e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(e0, e3))) & ((e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, e0))) & ((e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(e0, e3))) & ((e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, e0))) & ((e0 = op(e3, e3)) | (e0 = op(e2, e3)) | (e0 = op(e1, e3)) | (e0 = op(e0, e3))) & ((e0 = op(e3, e3)) | (e0 = op(e3, e2)) | (e0 = op(e3, e1)) | (e0 = op(e3, e0))) & ((e3 = op(e3, e2)) | (e3 = op(e2, e2)) | (e3 = op(e1, e2)) | (e3 = op(e0, e2))) & ((e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, e0))) & ((e2 = op(e3, e2)) | (e2 = op(e2, e2)) | (e2 = op(e1, e2)) | (e2 = op(e0, e2))) & ((e2 = op(e2, e3)) | (e2 = op(e2, e2)) | (e2 = op(e2, e1)) | (e2 = op(e2, e0))) & ((e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(e0, e2))) & ((e1 = op(e2, e3)) | (e1 = op(e2, e2)) | (e1 = op(e2, e1)) | (e1 = op(e2, e0))) & ((e0 = op(e3, e2)) | (e0 = op(e2, e2)) | (e0 = op(e1, e2)) | (e0 = op(e0, e2))) & ((e0 = op(e2, e3)) | (e0 = op(e2, e2)) | (e0 = op(e2, e1)) | (e0 = op(e2, e0))) & ((e3 = op(e3, e1)) | (e3 = op(e2, e1)) | (e3 = op(e1, e1)) | (e3 = op(e0, e1))) & ((e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, e0))) & ((e2 = op(e3, e1)) | (e2 = op(e2, e1)) | (e2 = op(e1, e1)) | (e2 = op(e0, e1))) & ((e2 = op(e1, e3)) | (e2 = op(e1, e2)) | (e2 = op(e1, e1)) | (e2 = op(e1, e0))) & ((e1 = op(e3, e1)) | (e1 = op(e2, e1)) | (e1 = op(e1, e1)) | (e1 = op(e0, e1))) & ((e1 = op(e1, e3)) | (e1 = op(e1, e2)) | (e1 = op(e1, e1)) | (e1 = op(e1, e0))) & ((e0 = op(e3, e1)) | (e0 = op(e2, e1)) | (e0 = op(e1, e1)) | (e0 = op(e0, e1))) & ((e0 = op(e1, e3)) | (e0 = op(e1, e2)) | (e0 = op(e1, e1)) | (e0 = op(e1, e0))) & ((e3 = op(e3, e0)) | (e3 = op(e2, e0)) | (e3 = op(e1, e0)) | (op(e0, e0) = e3)) & ((e3 = op(e0, e3)) | (e3 = op(e0, e2)) | (e3 = op(e0, e1)) | (op(e0, e0) = e3)) & ((e2 = op(e3, e0)) | (e2 = op(e2, e0)) | (e2 = op(e1, e0)) | (op(e0, e0) = e2)) & ((e2 = op(e0, e3)) | (e2 = op(e0, e2)) | (e2 = op(e0, e1)) | (op(e0, e0) = e2)) & ((e1 = op(e3, e0)) | (e1 = op(e2, e0)) | (e1 = op(e1, e0)) | (op(e0, e0) = e1)) & ((e1 = op(e0, e3)) | (e1 = op(e0, e2)) | (e1 = op(e0, e1)) | (op(e0, e0) = e1)) & ((e0 = op(e3, e0)) | (e0 = op(e2, e0)) | (e0 = op(e1, e0)) | (e0 = op(e0, e0))) & ((e0 = op(e0, e3)) | (e0 = op(e0, e2)) | (e0 = op(e0, e1)) | (e0 = op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG131+1.p', ax2)).
fof(f613, plain, (spl15_61 | spl15_45 | spl15_29 | spl15_13), inference(avatar_split_clause, [], [f102, f363, f431, f499, f567])).
fof(f102, plain, ((e0 = op(e3, e0)) | (e0 = op(e2, e0)) | (e0 = op(e1, e0)) | (e0 = op(e0, e0))), inference(cnf_transformation, [], [f2])).
fof(f612, plain, (spl15_62 | spl15_58 | spl15_54 | spl15_50), inference(avatar_split_clause, [], [f103, f520, f537, f554, f571])).
fof(f103, plain, ((e1 = op(e0, e3)) | (e1 = op(e0, e2)) | (e1 = op(e0, e1)) | (op(e0, e0) = e1)), inference(cnf_transformation, [], [f2])).
fof(f611, plain, (spl15_62 | spl15_46 | spl15_30 | spl15_14), inference(avatar_split_clause, [], [f104, f367, f435, f503, f571])).
fof(f104, plain, ((e1 = op(e3, e0)) | (e1 = op(e2, e0)) | (e1 = op(e1, e0)) | (op(e0, e0) = e1)), inference(cnf_transformation, [], [f2])).
fof(f610, plain, (spl15_63 | spl15_59 | spl15_55 | spl15_51), inference(avatar_split_clause, [], [f105, f524, f541, f558, f575])).
fof(f105, plain, ((e2 = op(e0, e3)) | (e2 = op(e0, e2)) | (e2 = op(e0, e1)) | (op(e0, e0) = e2)), inference(cnf_transformation, [], [f2])).
fof(f609, plain, (spl15_63 | spl15_47 | spl15_31 | spl15_15), inference(avatar_split_clause, [], [f106, f371, f439, f507, f575])).
fof(f106, plain, ((e2 = op(e3, e0)) | (e2 = op(e2, e0)) | (e2 = op(e1, e0)) | (op(e0, e0) = e2)), inference(cnf_transformation, [], [f2])).
fof(f608, plain, (spl15_64 | spl15_60 | spl15_56 | spl15_52), inference(avatar_split_clause, [], [f107, f528, f545, f562, f579])).
fof(f107, plain, ((e3 = op(e0, e3)) | (e3 = op(e0, e2)) | (e3 = op(e0, e1)) | (op(e0, e0) = e3)), inference(cnf_transformation, [], [f2])).
fof(f607, plain, (spl15_64 | spl15_48 | spl15_32 | spl15_16), inference(avatar_split_clause, [], [f108, f375, f443, f511, f579])).
fof(f108, plain, ((e3 = op(e3, e0)) | (e3 = op(e2, e0)) | (e3 = op(e1, e0)) | (op(e0, e0) = e3)), inference(cnf_transformation, [], [f2])).
fof(f606, plain, (spl15_45 | spl15_41 | spl15_37 | spl15_33), inference(avatar_split_clause, [], [f109, f448, f465, f482, f499])).
fof(f109, plain, ((e0 = op(e1, e3)) | (e0 = op(e1, e2)) | (e0 = op(e1, e1)) | (e0 = op(e1, e0))), inference(cnf_transformation, [], [f2])).
fof(f605, plain, (spl15_57 | spl15_41 | spl15_25 | spl15_9), inference(avatar_split_clause, [], [f110, f346, f414, f482, f550])).
fof(f110, plain, ((e0 = op(e3, e1)) | (e0 = op(e2, e1)) | (e0 = op(e1, e1)) | (e0 = op(e0, e1))), inference(cnf_transformation, [], [f2])).
fof(f604, plain, (spl15_46 | spl15_42 | spl15_38 | spl15_34), inference(avatar_split_clause, [], [f111, f452, f469, f486, f503])).
fof(f111, plain, ((e1 = op(e1, e3)) | (e1 = op(e1, e2)) | (e1 = op(e1, e1)) | (e1 = op(e1, e0))), inference(cnf_transformation, [], [f2])).
fof(f603, plain, (spl15_58 | spl15_42 | spl15_26 | spl15_10), inference(avatar_split_clause, [], [f112, f350, f418, f486, f554])).
fof(f112, plain, ((e1 = op(e3, e1)) | (e1 = op(e2, e1)) | (e1 = op(e1, e1)) | (e1 = op(e0, e1))), inference(cnf_transformation, [], [f2])).
fof(f602, plain, (spl15_47 | spl15_43 | spl15_39 | spl15_35), inference(avatar_split_clause, [], [f113, f456, f473, f490, f507])).
fof(f113, plain, ((e2 = op(e1, e3)) | (e2 = op(e1, e2)) | (e2 = op(e1, e1)) | (e2 = op(e1, e0))), inference(cnf_transformation, [], [f2])).
fof(f601, plain, (spl15_59 | spl15_43 | spl15_27 | spl15_11), inference(avatar_split_clause, [], [f114, f354, f422, f490, f558])).
fof(f114, plain, ((e2 = op(e3, e1)) | (e2 = op(e2, e1)) | (e2 = op(e1, e1)) | (e2 = op(e0, e1))), inference(cnf_transformation, [], [f2])).
fof(f600, plain, (spl15_48 | spl15_44 | spl15_40 | spl15_36), inference(avatar_split_clause, [], [f115, f460, f477, f494, f511])).
fof(f115, plain, ((e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, e0))), inference(cnf_transformation, [], [f2])).
fof(f599, plain, (spl15_60 | spl15_44 | spl15_28 | spl15_12), inference(avatar_split_clause, [], [f116, f358, f426, f494, f562])).
fof(f116, plain, ((e3 = op(e3, e1)) | (e3 = op(e2, e1)) | (e3 = op(e1, e1)) | (e3 = op(e0, e1))), inference(cnf_transformation, [], [f2])).
fof(f598, plain, (spl15_29 | spl15_25 | spl15_21 | spl15_17), inference(avatar_split_clause, [], [f117, f380, f397, f414, f431])).
fof(f117, plain, ((e0 = op(e2, e3)) | (e0 = op(e2, e2)) | (e0 = op(e2, e1)) | (e0 = op(e2, e0))), inference(cnf_transformation, [], [f2])).
fof(f597, plain, (spl15_53 | spl15_37 | spl15_21 | spl15_5), inference(avatar_split_clause, [], [f118, f329, f397, f465, f533])).
fof(f118, plain, ((e0 = op(e3, e2)) | (e0 = op(e2, e2)) | (e0 = op(e1, e2)) | (e0 = op(e0, e2))), inference(cnf_transformation, [], [f2])).
fof(f596, plain, (spl15_30 | spl15_26 | spl15_22 | spl15_18), inference(avatar_split_clause, [], [f119, f384, f401, f418, f435])).
fof(f119, plain, ((e1 = op(e2, e3)) | (e1 = op(e2, e2)) | (e1 = op(e2, e1)) | (e1 = op(e2, e0))), inference(cnf_transformation, [], [f2])).
fof(f595, plain, (spl15_54 | spl15_38 | spl15_22 | spl15_6), inference(avatar_split_clause, [], [f120, f333, f401, f469, f537])).
fof(f120, plain, ((e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(e0, e2))), inference(cnf_transformation, [], [f2])).
fof(f594, plain, (spl15_31 | spl15_27 | spl15_23 | spl15_19), inference(avatar_split_clause, [], [f121, f388, f405, f422, f439])).
fof(f121, plain, ((e2 = op(e2, e3)) | (e2 = op(e2, e2)) | (e2 = op(e2, e1)) | (e2 = op(e2, e0))), inference(cnf_transformation, [], [f2])).
fof(f593, plain, (spl15_55 | spl15_39 | spl15_23 | spl15_7), inference(avatar_split_clause, [], [f122, f337, f405, f473, f541])).
fof(f122, plain, ((e2 = op(e3, e2)) | (e2 = op(e2, e2)) | (e2 = op(e1, e2)) | (e2 = op(e0, e2))), inference(cnf_transformation, [], [f2])).
fof(f592, plain, (spl15_32 | spl15_28 | spl15_24 | spl15_20), inference(avatar_split_clause, [], [f123, f392, f409, f426, f443])).
fof(f123, plain, ((e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, e0))), inference(cnf_transformation, [], [f2])).
fof(f591, plain, (spl15_56 | spl15_40 | spl15_24 | spl15_8), inference(avatar_split_clause, [], [f124, f341, f409, f477, f545])).
fof(f124, plain, ((e3 = op(e3, e2)) | (e3 = op(e2, e2)) | (e3 = op(e1, e2)) | (e3 = op(e0, e2))), inference(cnf_transformation, [], [f2])).
fof(f590, plain, (spl15_13 | spl15_9 | spl15_5 | spl15_1), inference(avatar_split_clause, [], [f125, f312, f329, f346, f363])).
fof(f125, plain, ((e0 = op(e3, e3)) | (e0 = op(e3, e2)) | (e0 = op(e3, e1)) | (e0 = op(e3, e0))), inference(cnf_transformation, [], [f2])).
fof(f589, plain, (spl15_49 | spl15_33 | spl15_17 | spl15_1), inference(avatar_split_clause, [], [f126, f312, f380, f448, f516])).
fof(f126, plain, ((e0 = op(e3, e3)) | (e0 = op(e2, e3)) | (e0 = op(e1, e3)) | (e0 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f588, plain, (spl15_14 | spl15_10 | spl15_6 | spl15_2), inference(avatar_split_clause, [], [f127, f316, f333, f350, f367])).
fof(f127, plain, ((e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, e0))), inference(cnf_transformation, [], [f2])).
fof(f587, plain, (spl15_50 | spl15_34 | spl15_18 | spl15_2), inference(avatar_split_clause, [], [f128, f316, f384, f452, f520])).
fof(f128, plain, ((e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f586, plain, (spl15_15 | spl15_11 | spl15_7 | spl15_3), inference(avatar_split_clause, [], [f129, f320, f337, f354, f371])).
fof(f129, plain, ((e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, e0))), inference(cnf_transformation, [], [f2])).
fof(f585, plain, (spl15_51 | spl15_35 | spl15_19 | spl15_3), inference(avatar_split_clause, [], [f130, f320, f388, f456, f524])).
fof(f130, plain, ((e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f584, plain, (spl15_16 | spl15_12 | spl15_8 | spl15_4), inference(avatar_split_clause, [], [f131, f324, f341, f358, f375])).
fof(f131, plain, ((e3 = op(e3, e3)) | (e3 = op(e3, e2)) | (e3 = op(e3, e1)) | (e3 = op(e3, e0))), inference(cnf_transformation, [], [f2])).
fof(f583, plain, (spl15_52 | spl15_36 | spl15_20 | spl15_4), inference(avatar_split_clause, [], [f132, f324, f392, f460, f528])).
fof(f132, plain, ((e3 = op(e3, e3)) | (e3 = op(e2, e3)) | (e3 = op(e1, e3)) | (e3 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f582, plain, (spl15_61 | spl15_62 | spl15_63 | spl15_64), inference(avatar_split_clause, [], [f85, f579, f575, f571, f567])).
fof(f85, plain, ((op(e0, e0) = e3) | (op(e0, e0) = e2) | (op(e0, e0) = e1) | (e0 = op(e0, e0))), inference(cnf_transformation, [], [f1])).
fof(f1, plain, (((e3 = op(e3, e3)) | (e2 = op(e3, e3)) | (e1 = op(e3, e3)) | (e0 = op(e3, e3))) & ((e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (e0 = op(e3, e2))) & ((e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (e0 = op(e3, e1))) & ((e3 = op(e3, e0)) | (e2 = op(e3, e0)) | (e1 = op(e3, e0)) | (e0 = op(e3, e0))) & ((e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))) & ((e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (e0 = op(e2, e2))) & ((e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))) & ((e3 = op(e2, e0)) | (e2 = op(e2, e0)) | (e1 = op(e2, e0)) | (e0 = op(e2, e0))) & ((e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (e0 = op(e1, e3))) & ((e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (e0 = op(e1, e2))) & ((e3 = op(e1, e1)) | (e2 = op(e1, e1)) | (e1 = op(e1, e1)) | (e0 = op(e1, e1))) & ((e3 = op(e1, e0)) | (e2 = op(e1, e0)) | (e1 = op(e1, e0)) | (e0 = op(e1, e0))) & ((e3 = op(e0, e3)) | (e2 = op(e0, e3)) | (e1 = op(e0, e3)) | (e0 = op(e0, e3))) & ((e3 = op(e0, e2)) | (e2 = op(e0, e2)) | (e1 = op(e0, e2)) | (e0 = op(e0, e2))) & ((e3 = op(e0, e1)) | (e2 = op(e0, e1)) | (e1 = op(e0, e1)) | (e0 = op(e0, e1))) & ((op(e0, e0) = e3) | (op(e0, e0) = e2) | (op(e0, e0) = e1) | (e0 = op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG131+1.p', ax1)).
fof(f565, plain, (spl15_57 | spl15_58 | spl15_59 | spl15_60), inference(avatar_split_clause, [], [f86, f562, f558, f554, f550])).
fof(f86, plain, ((e3 = op(e0, e1)) | (e2 = op(e0, e1)) | (e1 = op(e0, e1)) | (e0 = op(e0, e1))), inference(cnf_transformation, [], [f1])).
fof(f548, plain, (spl15_53 | spl15_54 | spl15_55 | spl15_56), inference(avatar_split_clause, [], [f87, f545, f541, f537, f533])).
fof(f87, plain, ((e3 = op(e0, e2)) | (e2 = op(e0, e2)) | (e1 = op(e0, e2)) | (e0 = op(e0, e2))), inference(cnf_transformation, [], [f1])).
fof(f531, plain, (spl15_49 | spl15_50 | spl15_51 | spl15_52), inference(avatar_split_clause, [], [f88, f528, f524, f520, f516])).
fof(f88, plain, ((e3 = op(e0, e3)) | (e2 = op(e0, e3)) | (e1 = op(e0, e3)) | (e0 = op(e0, e3))), inference(cnf_transformation, [], [f1])).
fof(f514, plain, (spl15_45 | spl15_46 | spl15_47 | spl15_48), inference(avatar_split_clause, [], [f89, f511, f507, f503, f499])).
fof(f89, plain, ((e3 = op(e1, e0)) | (e2 = op(e1, e0)) | (e1 = op(e1, e0)) | (e0 = op(e1, e0))), inference(cnf_transformation, [], [f1])).
fof(f497, plain, (spl15_41 | spl15_42 | spl15_43 | spl15_44), inference(avatar_split_clause, [], [f90, f494, f490, f486, f482])).
fof(f90, plain, ((e3 = op(e1, e1)) | (e2 = op(e1, e1)) | (e1 = op(e1, e1)) | (e0 = op(e1, e1))), inference(cnf_transformation, [], [f1])).
fof(f480, plain, (spl15_37 | spl15_38 | spl15_39 | spl15_40), inference(avatar_split_clause, [], [f91, f477, f473, f469, f465])).
fof(f91, plain, ((e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (e0 = op(e1, e2))), inference(cnf_transformation, [], [f1])).
fof(f463, plain, (spl15_33 | spl15_34 | spl15_35 | spl15_36), inference(avatar_split_clause, [], [f92, f460, f456, f452, f448])).
fof(f92, plain, ((e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (e0 = op(e1, e3))), inference(cnf_transformation, [], [f1])).
fof(f446, plain, (spl15_29 | spl15_30 | spl15_31 | spl15_32), inference(avatar_split_clause, [], [f93, f443, f439, f435, f431])).
fof(f93, plain, ((e3 = op(e2, e0)) | (e2 = op(e2, e0)) | (e1 = op(e2, e0)) | (e0 = op(e2, e0))), inference(cnf_transformation, [], [f1])).
fof(f429, plain, (spl15_25 | spl15_26 | spl15_27 | spl15_28), inference(avatar_split_clause, [], [f94, f426, f422, f418, f414])).
fof(f94, plain, ((e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))), inference(cnf_transformation, [], [f1])).
fof(f412, plain, (spl15_21 | spl15_22 | spl15_23 | spl15_24), inference(avatar_split_clause, [], [f95, f409, f405, f401, f397])).
fof(f95, plain, ((e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (e0 = op(e2, e2))), inference(cnf_transformation, [], [f1])).
fof(f395, plain, (spl15_17 | spl15_18 | spl15_19 | spl15_20), inference(avatar_split_clause, [], [f96, f392, f388, f384, f380])).
fof(f96, plain, ((e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))), inference(cnf_transformation, [], [f1])).
fof(f378, plain, (spl15_13 | spl15_14 | spl15_15 | spl15_16), inference(avatar_split_clause, [], [f97, f375, f371, f367, f363])).
fof(f97, plain, ((e3 = op(e3, e0)) | (e2 = op(e3, e0)) | (e1 = op(e3, e0)) | (e0 = op(e3, e0))), inference(cnf_transformation, [], [f1])).
fof(f361, plain, (spl15_9 | spl15_10 | spl15_11 | spl15_12), inference(avatar_split_clause, [], [f98, f358, f354, f350, f346])).
fof(f98, plain, ((e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (e0 = op(e3, e1))), inference(cnf_transformation, [], [f1])).
fof(f344, plain, (spl15_5 | spl15_6 | spl15_7 | spl15_8), inference(avatar_split_clause, [], [f99, f341, f337, f333, f329])).
fof(f99, plain, ((e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (e0 = op(e3, e2))), inference(cnf_transformation, [], [f1])).
fof(f327, plain, (spl15_1 | spl15_2 | spl15_3 | spl15_4), inference(avatar_split_clause, [], [f100, f324, f320, f316, f312])).
fof(f100, plain, ((e3 = op(e3, e3)) | (e2 = op(e3, e3)) | (e1 = op(e3, e3)) | (e0 = op(e3, e3))), inference(cnf_transformation, [], [f1])).