fof(f3742, plain, $false, inference(avatar_sat_refutation, [], [f449, f466, f483, f500, f517, f534, f551, f568, f585, f602, f619, f636, f653, f670, f687, f704, f705, f706, f707, f708, f709, f710, f711, f712, f713, f714, f715, f716, f717, f718, f719, f720, f721, f722, f723, f724, f725, f726, f727, f728, f729, f730, f731, f732, f733, f734, f735, f736, f745, f750, f755, f760, f761, f762, f767, f768, f769, f770, f771, f772, f778, f779, f780, f781, f782, f791, f796, f801, f806, f807, f808, f814, f818, f823, f825, f826, f827, f828, f833, f834, f835, f836, f837, f838, f847, f852, f857, f862, f863, f864, f869, f870, f871, f872, f873, f874, f881, f884, f889, f890, f891, f892, f893, f894, f903, f908, f913, f918, f919, f920, f925, f926, f927, f928, f929, f930, f935, f936, f937, f939, f940, f948, f950, f955, f956, f957, f958, f959, f960, f965, f966, f967, f968, f969, f970, f975, f976, f977, f978, f979, f980, f985, f986, f987, f988, f989, f990, f995, f997, f999, f1000, f1005, f1006, f1007, f1008, f1009, f1010, f1015, f1016, f1017, f1018, f1019, f1020, f1025, f1026, f1027, f1028, f1029, f1030, f1035, f1036, f1037, f1038, f1039, f1040, f1046, f1049, f1050, f1055, f1056, f1057, f1058, f1059, f1060, f1065, f1066, f1067, f1068, f1069, f1070, f1075, f1076, f1077, f1078, f1079, f1080, f1085, f1086, f1087, f1088, f1089, f1090, f1095, f1096, f1099, f1100, f1101, f1102, f1104, f1106, f1108, f1109, f1110, f1111, f1112, f1131, f1136, f1141, f1142, f1151, f1156, f1161, f1172, f1181, f1186, f1196, f1211, f1231, f1232, f1244, f1245, f1247, f1248, f1250, f1257, f1268, f1270, f1271, f1277, f1281, f1284, f1287, f1294, f1299, f1310, f1312, f1313, f1324, f1326, f1331, f1333, f1336, f1342, f1343, f1347, f1353, f1370, f1372, f1375, f1378, f1381, f1389, f1392, f1397, f1401, f1404, f1407, f1410, f1411, f1414, f1417, f1423, f1427, f1435, f1443, f1448, f1480, f1482, f1485, f1491, f1503, f1504, f1505, f1517, f1528, f1546, f1555, f1565, f1568, f1574, f1576, f1579, f1586, f1588, f1598, f1613, f1628, f1630, f1633, f1635, f1639, f1668, f1679, f1697, f1725, f1745, f1757, f1764, f1767, f1772, f1775, f1781, f1786, f1790, f1797, f1801, f1802, f1811, f1821, f1836, f1839, f1842, f1853, f1856, f1898, f1904, f1946, f1949, f1952, f1956, f1966, f1975, f1976, f1982, f1998, f2014, f2018, f2055, f2064, f2067, f2071, f2079, f2099, f2103, f2113, f2117, f2135, f2136, f2140, f2155, f2186, f2209, f2214, f2229, f2234, f2248, f2249, f2267, f2278, f2283, f2290, f2304, f2312, f2314, f2324, f2340, f2355, f2359, f2365, f2370, f2380, f2384, f2396, f2397, f2401, f2431, f2433, f2437, f2442, f2461, f2467, f2482, f2494, f2497, f2527, f2536, f2573, f2584, f2594, f2627, f2628, f2632, f2636, f2653, f2655, f2663, f2668, f2670, f2674, f2679, f2699, f2700, f2705, f2712, f2728, f2740, f2743, f2747, f2752, f2758, f2767, f2776, f2786, f2787, f2804, f2815, f2818, f2844, f2845, f2846, f2847, f2849, f2856, f2871, f2879, f2880, f2893, f2897, f2904, f2913, f2917, f2933, f2962, f2964, f2979, f2990, f2999, f3009, f3034, f3054, f3069, f3081, f3103, f3108, f3110, f3112, f3115, f3117, f3141, f3146, f3151, f3158, f3167, f3168, f3175, f3179, f3183, f3192, f3208, f3231, f3232, f3233, f3244, f3245, f3251, f3262, f3263, f3272, f3273, f3279, f3285, f3292, f3299, f3300, f3307, f3309, f3311, f3322, f3329, f3361, f3375, f3376, f3379, f3388, f3389, f3390, f3395, f3396, f3402, f3409, f3410, f3420, f3443, f3459, f3460, f3462, f3470, f3471, f3477, f3506, f3507, f3514, f3523, f3524, f3530, f3535, f3543, f3549, f3550, f3553, f3563, f3570, f3581, f3591, f3596, f3621, f3622, f3626, f3636, f3637, f3644, f3645, f3650, f3657, f3658, f3668, f3676, f3685, f3697, f3702, f3709, f3711, f3721, f3741])).
fof(f3741, plain, (~ spl30_2 | ~ spl30_47 | ~ spl30_129 | spl30_130), inference(avatar_contradiction_clause, [], [f3740])).
fof(f3740, plain, ($false | (~ spl30_2 | ~ spl30_47 | ~ spl30_129 | spl30_130)), inference(subsumption_resolution, [], [f3737, f631])).
fof(f631, plain, ((e2 = op(e1, e0)) | ~ spl30_47), inference(avatar_component_clause, [], [f629])).
fof(f629, plain, (spl30_47 <=> (e2 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_47])])).
fof(f3737, plain, (~ (e2 = op(e1, e0)) | (~ spl30_2 | ~ spl30_129 | spl30_130)), inference(backward_demodulation, [], [f3601, f440])).
fof(f440, plain, ((e1 = op(e3, e3)) | ~ spl30_2), inference(avatar_component_clause, [], [f438])).
fof(f438, plain, (spl30_2 <=> (e1 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_2])])).
fof(f3601, plain, (~ (e2 = op(op(e3, e3), e0)) | (~ spl30_129 | spl30_130)), inference(backward_demodulation, [], [f1210, f1205])).
fof(f1205, plain, ((e0 = op(op(e3, e3), e3)) | ~ spl30_129), inference(avatar_component_clause, [], [f1204])).
fof(f1204, plain, (spl30_129 <=> (e0 = op(op(e3, e3), e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_129])])).
fof(f1210, plain, (~ (e2 = op(op(e3, e3), op(op(e3, e3), e3))) | spl30_130), inference(avatar_component_clause, [], [f1208])).
fof(f1208, plain, (spl30_130 <=> (e2 = op(op(e3, e3), op(op(e3, e3), e3)))), introduced(avatar_definition, [new_symbols(naming, [spl30_130])])).
fof(f3721, plain, (~ spl30_4 | ~ spl30_16), inference(avatar_split_clause, [], [f3719, f497, f446])).
fof(f446, plain, (spl30_4 <=> (e3 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_4])])).
fof(f497, plain, (spl30_16 <=> (e3 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_16])])).
fof(f3719, plain, (~ (e3 = op(e3, e3)) | ~ spl30_16), inference(backward_demodulation, [], [f207, f499])).
fof(f499, plain, ((e3 = op(e3, e0)) | ~ spl30_16), inference(avatar_component_clause, [], [f497])).
fof(f207, plain, ~ (op(e3, e0) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f3, plain, (~ (op(e3, e2) = op(e3, e3)) & ~ (op(e3, e1) = op(e3, e3)) & ~ (op(e3, e1) = op(e3, e2)) & ~ (op(e3, e0) = op(e3, e3)) & ~ (op(e3, e0) = op(e3, e2)) & ~ (op(e3, e0) = op(e3, e1)) & ~ (op(e2, e2) = op(e2, e3)) & ~ (op(e2, e1) = op(e2, e3)) & ~ (op(e2, e1) = op(e2, e2)) & ~ (op(e2, e0) = op(e2, e3)) & ~ (op(e2, e0) = op(e2, e2)) & ~ (op(e2, e0) = op(e2, e1)) & ~ (op(e1, e2) = op(e1, e3)) & ~ (op(e1, e1) = op(e1, e3)) & ~ (op(e1, e1) = op(e1, e2)) & ~ (op(e1, e0) = op(e1, e3)) & ~ (op(e1, e0) = op(e1, e2)) & ~ (op(e1, e0) = op(e1, e1)) & ~ (op(e0, e2) = op(e0, e3)) & ~ (op(e0, e1) = op(e0, e3)) & ~ (op(e0, e1) = op(e0, e2)) & ~ (op(e0, e0) = op(e0, e3)) & ~ (op(e0, e0) = op(e0, e2)) & ~ (op(e0, e0) = op(e0, e1)) & ~ (op(e2, e3) = op(e3, e3)) & ~ (op(e1, e3) = op(e3, e3)) & ~ (op(e1, e3) = op(e2, e3)) & ~ (op(e0, e3) = op(e3, e3)) & ~ (op(e0, e3) = op(e2, e3)) & ~ (op(e0, e3) = op(e1, e3)) & ~ (op(e2, e2) = op(e3, e2)) & ~ (op(e1, e2) = op(e3, e2)) & ~ (op(e1, e2) = op(e2, e2)) & ~ (op(e0, e2) = op(e3, e2)) & ~ (op(e0, e2) = op(e2, e2)) & ~ (op(e0, e2) = op(e1, e2)) & ~ (op(e2, e1) = op(e3, e1)) & ~ (op(e1, e1) = op(e3, e1)) & ~ (op(e1, e1) = op(e2, e1)) & ~ (op(e0, e1) = op(e3, e1)) & ~ (op(e0, e1) = op(e2, e1)) & ~ (op(e0, e1) = op(e1, e1)) & ~ (op(e2, e0) = op(e3, e0)) & ~ (op(e1, e0) = op(e3, e0)) & ~ (op(e1, e0) = op(e2, e0)) & ~ (op(e0, e0) = op(e3, e0)) & ~ (op(e0, e0) = op(e2, e0)) & ~ (op(e0, e0) = op(e1, e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG159+1.p', ax3)).
fof(f3711, plain, (~ spl30_22 | spl30_38 | ~ spl30_126), inference(avatar_contradiction_clause, [], [f3710])).
fof(f3710, plain, ($false | (~ spl30_22 | spl30_38 | ~ spl30_126)), inference(subsumption_resolution, [], [f3707, f592])).
fof(f592, plain, (~ (e1 = op(e1, e2)) | spl30_38), inference(avatar_component_clause, [], [f591])).
fof(f591, plain, (spl30_38 <=> (e1 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_38])])).
fof(f3707, plain, ((e1 = op(e1, e2)) | (~ spl30_22 | ~ spl30_126)), inference(backward_demodulation, [], [f1189, f525])).
fof(f525, plain, ((e1 = op(e2, e2)) | ~ spl30_22), inference(avatar_component_clause, [], [f523])).
fof(f523, plain, (spl30_22 <=> (e1 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_22])])).
fof(f1189, plain, ((e1 = op(op(e2, e2), e2)) | ~ spl30_126), inference(avatar_component_clause, [], [f1188])).
fof(f1188, plain, (spl30_126 <=> (e1 = op(op(e2, e2), e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_126])])).
fof(f3709, plain, (~ spl30_6 | ~ spl30_22), inference(avatar_split_clause, [], [f3703, f523, f455])).
fof(f455, plain, (spl30_6 <=> (e1 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_6])])).
fof(f3703, plain, (~ (e1 = op(e3, e2)) | ~ spl30_22), inference(backward_demodulation, [], [f180, f525])).
fof(f180, plain, ~ (op(e2, e2) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f3702, plain, (~ spl30_24 | ~ spl30_28), inference(avatar_split_clause, [], [f3699, f548, f531])).
fof(f531, plain, (spl30_24 <=> (e3 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_24])])).
fof(f548, plain, (spl30_28 <=> (e3 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_28])])).
fof(f3699, plain, (~ (e3 = op(e2, e2)) | ~ spl30_28), inference(backward_demodulation, [], [f202, f550])).
fof(f550, plain, ((e3 = op(e2, e1)) | ~ spl30_28), inference(avatar_component_clause, [], [f548])).
fof(f202, plain, ~ (op(e2, e1) = op(e2, e2)), inference(cnf_transformation, [], [f3])).
fof(f3697, plain, (~ spl30_17 | ~ spl30_29), inference(avatar_split_clause, [], [f3692, f553, f502])).
fof(f502, plain, (spl30_17 <=> (e0 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_17])])).
fof(f553, plain, (spl30_29 <=> (e0 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_29])])).
fof(f3692, plain, (~ (e0 = op(e2, e3)) | ~ spl30_29), inference(backward_demodulation, [], [f201, f555])).
fof(f555, plain, ((e0 = op(e2, e0)) | ~ spl30_29), inference(avatar_component_clause, [], [f553])).
fof(f201, plain, ~ (op(e2, e0) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f3685, plain, (~ spl30_17 | ~ spl30_33), inference(avatar_split_clause, [], [f3679, f570, f502])).
fof(f570, plain, (spl30_33 <=> (e0 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_33])])).
fof(f3679, plain, (~ (e0 = op(e2, e3)) | ~ spl30_33), inference(backward_demodulation, [], [f184, f572])).
fof(f572, plain, ((e0 = op(e1, e3)) | ~ spl30_33), inference(avatar_component_clause, [], [f570])).
fof(f184, plain, ~ (op(e1, e3) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f3676, plain, (~ spl30_24 | ~ spl30_40), inference(avatar_split_clause, [], [f3672, f599, f531])).
fof(f599, plain, (spl30_40 <=> (e3 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_40])])).
fof(f3672, plain, (~ (e3 = op(e2, e2)) | ~ spl30_40), inference(backward_demodulation, [], [f178, f601])).
fof(f601, plain, ((e3 = op(e1, e2)) | ~ spl30_40), inference(avatar_component_clause, [], [f599])).
fof(f178, plain, ~ (op(e1, e2) = op(e2, e2)), inference(cnf_transformation, [], [f3])).
fof(f3668, plain, (~ spl30_34 | ~ spl30_42), inference(avatar_split_clause, [], [f3662, f608, f574])).
fof(f574, plain, (spl30_34 <=> (e1 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_34])])).
fof(f608, plain, (spl30_42 <=> (e1 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_42])])).
fof(f3662, plain, (~ (e1 = op(e1, e3)) | ~ spl30_42), inference(backward_demodulation, [], [f197, f610])).
fof(f610, plain, ((e1 = op(e1, e1)) | ~ spl30_42), inference(avatar_component_clause, [], [f608])).
fof(f197, plain, ~ (op(e1, e1) = op(e1, e3)), inference(cnf_transformation, [], [f3])).
fof(f3658, plain, (~ spl30_39 | ~ spl30_47), inference(avatar_split_clause, [], [f3654, f629, f595])).
fof(f595, plain, (spl30_39 <=> (e2 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_39])])).
fof(f3654, plain, (~ (e2 = op(e1, e2)) | ~ spl30_47), inference(backward_demodulation, [], [f194, f631])).
fof(f194, plain, ~ (op(e1, e0) = op(e1, e2)), inference(cnf_transformation, [], [f3])).
fof(f3657, plain, (~ spl30_15 | ~ spl30_47), inference(avatar_split_clause, [], [f3652, f629, f493])).
fof(f493, plain, (spl30_15 <=> (e2 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_15])])).
fof(f3652, plain, (~ (e2 = op(e3, e0)) | ~ spl30_47), inference(backward_demodulation, [], [f167, f631])).
fof(f167, plain, ~ (op(e1, e0) = op(e3, e0)), inference(cnf_transformation, [], [f3])).
fof(f3650, plain, (~ spl30_4 | ~ spl30_52), inference(avatar_split_clause, [], [f3648, f650, f446])).
fof(f650, plain, (spl30_52 <=> (e3 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_52])])).
fof(f3648, plain, (~ (e3 = op(e3, e3)) | ~ spl30_52), inference(backward_demodulation, [], [f183, f652])).
fof(f652, plain, ((e3 = op(e0, e3)) | ~ spl30_52), inference(avatar_component_clause, [], [f650])).
fof(f183, plain, ~ (op(e0, e3) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f3645, plain, (~ spl30_51 | ~ spl30_55), inference(avatar_split_clause, [], [f3641, f663, f646])).
fof(f646, plain, (spl30_51 <=> (e2 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_51])])).
fof(f663, plain, (spl30_55 <=> (e2 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_55])])).
fof(f3641, plain, (~ (e2 = op(e0, e3)) | ~ spl30_55), inference(backward_demodulation, [], [f192, f665])).
fof(f665, plain, ((e2 = op(e0, e2)) | ~ spl30_55), inference(avatar_component_clause, [], [f663])).
fof(f192, plain, ~ (op(e0, e2) = op(e0, e3)), inference(cnf_transformation, [], [f3])).
fof(f3644, plain, (~ spl30_39 | ~ spl30_55), inference(avatar_split_clause, [], [f3638, f663, f595])).
fof(f3638, plain, (~ (e2 = op(e1, e2)) | ~ spl30_55), inference(backward_demodulation, [], [f175, f665])).
fof(f175, plain, ~ (op(e0, e2) = op(e1, e2)), inference(cnf_transformation, [], [f3])).
fof(f3637, plain, (~ spl30_53 | ~ spl30_57), inference(avatar_split_clause, [], [f3630, f672, f655])).
fof(f655, plain, (spl30_53 <=> (e0 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_53])])).
fof(f672, plain, (spl30_57 <=> (e0 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_57])])).
fof(f3630, plain, (~ (e0 = op(e0, e2)) | ~ spl30_57), inference(backward_demodulation, [], [f190, f674])).
fof(f674, plain, ((e0 = op(e0, e1)) | ~ spl30_57), inference(avatar_component_clause, [], [f672])).
fof(f190, plain, ~ (op(e0, e1) = op(e0, e2)), inference(cnf_transformation, [], [f3])).
fof(f3636, plain, (~ spl30_9 | ~ spl30_57), inference(avatar_split_clause, [], [f3629, f672, f468])).
fof(f468, plain, (spl30_9 <=> (e0 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_9])])).
fof(f3629, plain, (~ (e0 = op(e3, e1)) | ~ spl30_57), inference(backward_demodulation, [], [f171, f674])).
fof(f171, plain, ~ (op(e0, e1) = op(e3, e1)), inference(cnf_transformation, [], [f3])).
fof(f3626, plain, (~ spl30_40 | ~ spl30_62 | spl30_112 | ~ spl30_114), inference(avatar_split_clause, [], [f3620, f1128, f1118, f693, f599])).
fof(f693, plain, (spl30_62 <=> (op(e0, e0) = e1)), introduced(avatar_definition, [new_symbols(naming, [spl30_62])])).
fof(f1118, plain, (spl30_112 <=> (e3 = op(op(e0, e0), op(op(e0, e0), e0)))), introduced(avatar_definition, [new_symbols(naming, [spl30_112])])).
fof(f1128, plain, (spl30_114 <=> (e2 = op(op(e0, e0), e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_114])])).
fof(f3620, plain, (~ (e3 = op(e1, e2)) | (~ spl30_62 | spl30_112 | ~ spl30_114)), inference(backward_demodulation, [], [f3606, f695])).
fof(f695, plain, ((op(e0, e0) = e1) | ~ spl30_62), inference(avatar_component_clause, [], [f693])).
fof(f3606, plain, (~ (e3 = op(op(e0, e0), e2)) | (spl30_112 | ~ spl30_114)), inference(forward_demodulation, [], [f1120, f1129])).
fof(f1129, plain, ((e2 = op(op(e0, e0), e0)) | ~ spl30_114), inference(avatar_component_clause, [], [f1128])).
fof(f1120, plain, (~ (e3 = op(op(e0, e0), op(op(e0, e0), e0))) | spl30_112), inference(avatar_component_clause, [], [f1118])).
fof(f3622, plain, (~ spl30_58 | ~ spl30_62), inference(avatar_split_clause, [], [f3610, f693, f676])).
fof(f676, plain, (spl30_58 <=> (e1 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_58])])).
fof(f3610, plain, (~ (e1 = op(e0, e1)) | ~ spl30_62), inference(backward_demodulation, [], [f187, f695])).
fof(f187, plain, ~ (op(e0, e0) = op(e0, e1)), inference(cnf_transformation, [], [f3])).
fof(f3621, plain, (~ spl30_30 | ~ spl30_62), inference(avatar_split_clause, [], [f3608, f693, f557])).
fof(f557, plain, (spl30_30 <=> (e1 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_30])])).
fof(f3608, plain, (~ (e1 = op(e2, e0)) | ~ spl30_62), inference(backward_demodulation, [], [f164, f695])).
fof(f164, plain, ~ (op(e0, e0) = op(e2, e0)), inference(cnf_transformation, [], [f3])).
fof(f3596, plain, (~ spl30_15 | ~ spl30_64 | spl30_114), inference(avatar_contradiction_clause, [], [f3595])).
fof(f3595, plain, ($false | (~ spl30_15 | ~ spl30_64 | spl30_114)), inference(subsumption_resolution, [], [f3594, f495])).
fof(f495, plain, ((e2 = op(e3, e0)) | ~ spl30_15), inference(avatar_component_clause, [], [f493])).
fof(f3594, plain, (~ (e2 = op(e3, e0)) | (~ spl30_64 | spl30_114)), inference(forward_demodulation, [], [f1130, f703])).
fof(f703, plain, ((op(e0, e0) = e3) | ~ spl30_64), inference(avatar_component_clause, [], [f701])).
fof(f701, plain, (spl30_64 <=> (op(e0, e0) = e3)), introduced(avatar_definition, [new_symbols(naming, [spl30_64])])).
fof(f1130, plain, (~ (e2 = op(op(e0, e0), e0)) | spl30_114), inference(avatar_component_clause, [], [f1128])).
fof(f3591, plain, (~ spl30_9 | ~ spl30_44 | spl30_117), inference(avatar_contradiction_clause, [], [f3590])).
fof(f3590, plain, ($false | (~ spl30_9 | ~ spl30_44 | spl30_117)), inference(subsumption_resolution, [], [f3589, f470])).
fof(f470, plain, ((e0 = op(e3, e1)) | ~ spl30_9), inference(avatar_component_clause, [], [f468])).
fof(f3589, plain, (~ (e0 = op(e3, e1)) | (~ spl30_44 | spl30_117)), inference(forward_demodulation, [], [f1146, f618])).
fof(f618, plain, ((e3 = op(e1, e1)) | ~ spl30_44), inference(avatar_component_clause, [], [f616])).
fof(f616, plain, (spl30_44 <=> (e3 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_44])])).
fof(f1146, plain, (~ (e0 = op(op(e1, e1), e1)) | spl30_117), inference(avatar_component_clause, [], [f1144])).
fof(f1144, plain, (spl30_117 <=> (e0 = op(op(e1, e1), e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_117])])).
fof(f3581, plain, (~ spl30_6 | ~ spl30_9 | ~ spl30_24 | spl30_127), inference(avatar_contradiction_clause, [], [f3580])).
fof(f3580, plain, ($false | (~ spl30_6 | ~ spl30_9 | ~ spl30_24 | spl30_127)), inference(subsumption_resolution, [], [f3579, f470])).
fof(f3579, plain, (~ (e0 = op(e3, e1)) | (~ spl30_6 | ~ spl30_24 | spl30_127)), inference(forward_demodulation, [], [f3578, f457])).
fof(f457, plain, ((e1 = op(e3, e2)) | ~ spl30_6), inference(avatar_component_clause, [], [f455])).
fof(f3578, plain, (~ (e0 = op(e3, op(e3, e2))) | (~ spl30_24 | spl30_127)), inference(forward_demodulation, [], [f1195, f533])).
fof(f533, plain, ((e3 = op(e2, e2)) | ~ spl30_24), inference(avatar_component_clause, [], [f531])).
fof(f1195, plain, (~ (e0 = op(op(e2, e2), op(op(e2, e2), e2))) | spl30_127), inference(avatar_component_clause, [], [f1193])).
fof(f1193, plain, (spl30_127 <=> (e0 = op(op(e2, e2), op(op(e2, e2), e2)))), introduced(avatar_definition, [new_symbols(naming, [spl30_127])])).
fof(f3570, plain, (~ spl30_19 | ~ spl30_24 | spl30_74), inference(avatar_split_clause, [], [f3565, f793, f531, f510])).
fof(f510, plain, (spl30_19 <=> (e2 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_19])])).
fof(f793, plain, (spl30_74 <=> (e2 = op(e2, op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl30_74])])).
fof(f3565, plain, (~ (e2 = op(e2, e3)) | (~ spl30_24 | spl30_74)), inference(backward_demodulation, [], [f795, f533])).
fof(f795, plain, (~ (e2 = op(e2, op(e2, e2))) | spl30_74), inference(avatar_component_clause, [], [f793])).
fof(f3563, plain, (~ spl30_19 | ~ spl30_27), inference(avatar_split_clause, [], [f3560, f544, f510])).
fof(f544, plain, (spl30_27 <=> (e2 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_27])])).
fof(f3560, plain, (~ (e2 = op(e2, e3)) | ~ spl30_27), inference(backward_demodulation, [], [f203, f546])).
fof(f546, plain, ((e2 = op(e2, e1)) | ~ spl30_27), inference(avatar_component_clause, [], [f544])).
fof(f203, plain, ~ (op(e2, e1) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f3553, plain, (~ spl30_9 | ~ spl30_15 | ~ spl30_44 | spl30_119), inference(avatar_contradiction_clause, [], [f3552])).
fof(f3552, plain, ($false | (~ spl30_9 | ~ spl30_15 | ~ spl30_44 | spl30_119)), inference(subsumption_resolution, [], [f3551, f495])).
fof(f3551, plain, (~ (e2 = op(e3, e0)) | (~ spl30_9 | ~ spl30_44 | spl30_119)), inference(forward_demodulation, [], [f3547, f470])).
fof(f3547, plain, (~ (e2 = op(e3, op(e3, e1))) | (~ spl30_44 | spl30_119)), inference(backward_demodulation, [], [f1155, f618])).
fof(f1155, plain, (~ (e2 = op(op(e1, e1), op(op(e1, e1), e1))) | spl30_119), inference(avatar_component_clause, [], [f1153])).
fof(f1153, plain, (spl30_119 <=> (e2 = op(op(e1, e1), op(op(e1, e1), e1)))), introduced(avatar_definition, [new_symbols(naming, [spl30_119])])).
fof(f3550, plain, (spl30_34 | ~ spl30_44 | ~ spl30_83), inference(avatar_split_clause, [], [f3546, f854, f616, f574])).
fof(f854, plain, (spl30_83 <=> (e1 = op(e1, op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl30_83])])).
fof(f3546, plain, ((e1 = op(e1, e3)) | (~ spl30_44 | ~ spl30_83)), inference(backward_demodulation, [], [f855, f618])).
fof(f855, plain, ((e1 = op(e1, op(e1, e1))) | ~ spl30_83), inference(avatar_component_clause, [], [f854])).
fof(f3549, plain, (~ spl30_28 | ~ spl30_44), inference(avatar_split_clause, [], [f3544, f616, f548])).
fof(f3544, plain, (~ (e3 = op(e2, e1)) | ~ spl30_44), inference(backward_demodulation, [], [f172, f618])).
fof(f172, plain, ~ (op(e1, e1) = op(e2, e1)), inference(cnf_transformation, [], [f3])).
fof(f3543, plain, (~ spl30_33 | ~ spl30_45), inference(avatar_split_clause, [], [f3538, f621, f570])).
fof(f621, plain, (spl30_45 <=> (e0 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_45])])).
fof(f3538, plain, (~ (e0 = op(e1, e3)) | ~ spl30_45), inference(backward_demodulation, [], [f195, f623])).
fof(f623, plain, ((e0 = op(e1, e0)) | ~ spl30_45), inference(avatar_component_clause, [], [f621])).
fof(f195, plain, ~ (op(e1, e0) = op(e1, e3)), inference(cnf_transformation, [], [f3])).
fof(f3535, plain, (~ spl30_19 | ~ spl30_51), inference(avatar_split_clause, [], [f3532, f646, f510])).
fof(f3532, plain, (~ (e2 = op(e2, e3)) | ~ spl30_51), inference(backward_demodulation, [], [f182, f648])).
fof(f648, plain, ((e2 = op(e0, e3)) | ~ spl30_51), inference(avatar_component_clause, [], [f646])).
fof(f182, plain, ~ (op(e0, e3) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f3530, plain, (~ spl30_21 | ~ spl30_53), inference(avatar_split_clause, [], [f3525, f655, f519])).
fof(f519, plain, (spl30_21 <=> (e0 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_21])])).
fof(f3525, plain, (~ (e0 = op(e2, e2)) | ~ spl30_53), inference(backward_demodulation, [], [f176, f657])).
fof(f657, plain, ((e0 = op(e0, e2)) | ~ spl30_53), inference(avatar_component_clause, [], [f655])).
fof(f176, plain, ~ (op(e0, e2) = op(e2, e2)), inference(cnf_transformation, [], [f3])).
fof(f3524, plain, (~ spl30_50 | ~ spl30_58), inference(avatar_split_clause, [], [f3519, f676, f642])).
fof(f642, plain, (spl30_50 <=> (e1 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_50])])).
fof(f3519, plain, (~ (e1 = op(e0, e3)) | ~ spl30_58), inference(backward_demodulation, [], [f191, f678])).
fof(f678, plain, ((e1 = op(e0, e1)) | ~ spl30_58), inference(avatar_component_clause, [], [f676])).
fof(f191, plain, ~ (op(e0, e1) = op(e0, e3)), inference(cnf_transformation, [], [f3])).
fof(f3523, plain, (~ spl30_42 | ~ spl30_58), inference(avatar_split_clause, [], [f3516, f676, f608])).
fof(f3516, plain, (~ (e1 = op(e1, e1)) | ~ spl30_58), inference(backward_demodulation, [], [f169, f678])).
fof(f169, plain, ~ (op(e0, e1) = op(e1, e1)), inference(cnf_transformation, [], [f3])).
fof(f3514, plain, (~ spl30_6 | ~ spl30_15 | ~ spl30_64 | spl30_115), inference(avatar_contradiction_clause, [], [f3513])).
fof(f3513, plain, ($false | (~ spl30_6 | ~ spl30_15 | ~ spl30_64 | spl30_115)), inference(subsumption_resolution, [], [f3512, f457])).
fof(f3512, plain, (~ (e1 = op(e3, e2)) | (~ spl30_15 | ~ spl30_64 | spl30_115)), inference(forward_demodulation, [], [f3504, f495])).
fof(f3504, plain, (~ (e1 = op(e3, op(e3, e0))) | (~ spl30_64 | spl30_115)), inference(backward_demodulation, [], [f1135, f703])).
fof(f1135, plain, (~ (e1 = op(op(e0, e0), op(op(e0, e0), e0))) | spl30_115), inference(avatar_component_clause, [], [f1133])).
fof(f1133, plain, (spl30_115 <=> (e1 = op(op(e0, e0), op(op(e0, e0), e0)))), introduced(avatar_definition, [new_symbols(naming, [spl30_115])])).
fof(f3507, plain, (~ spl30_56 | ~ spl30_64), inference(avatar_split_clause, [], [f3499, f701, f667])).
fof(f667, plain, (spl30_56 <=> (e3 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_56])])).
fof(f3499, plain, (~ (e3 = op(e0, e2)) | ~ spl30_64), inference(backward_demodulation, [], [f188, f703])).
fof(f188, plain, ~ (op(e0, e0) = op(e0, e2)), inference(cnf_transformation, [], [f3])).
fof(f3506, plain, (~ spl30_48 | ~ spl30_64), inference(avatar_split_clause, [], [f3497, f701, f633])).
fof(f633, plain, (spl30_48 <=> (e3 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_48])])).
fof(f3497, plain, (~ (e3 = op(e1, e0)) | ~ spl30_64), inference(backward_demodulation, [], [f163, f703])).
fof(f163, plain, ~ (op(e0, e0) = op(e1, e0)), inference(cnf_transformation, [], [f3])).
fof(f3477, plain, (~ spl30_5 | ~ spl30_9), inference(avatar_split_clause, [], [f3472, f468, f451])).
fof(f451, plain, (spl30_5 <=> (e0 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_5])])).
fof(f3472, plain, (~ (e0 = op(e3, e2)) | ~ spl30_9), inference(backward_demodulation, [], [f208, f470])).
fof(f208, plain, ~ (op(e3, e1) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f3471, plain, (~ spl30_5 | ~ spl30_15 | spl30_69), inference(avatar_split_clause, [], [f3469, f757, f493, f451])).
fof(f757, plain, (spl30_69 <=> (e0 = op(e3, op(e3, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl30_69])])).
fof(f3469, plain, (~ (e0 = op(e3, e2)) | (~ spl30_15 | spl30_69)), inference(backward_demodulation, [], [f759, f495])).
fof(f759, plain, (~ (e0 = op(e3, op(e3, e0))) | spl30_69), inference(avatar_component_clause, [], [f757])).
fof(f3470, plain, (~ spl30_11 | ~ spl30_15), inference(avatar_split_clause, [], [f3466, f493, f476])).
fof(f476, plain, (spl30_11 <=> (e2 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_11])])).
fof(f3466, plain, (~ (e2 = op(e3, e1)) | ~ spl30_15), inference(backward_demodulation, [], [f205, f495])).
fof(f205, plain, ~ (op(e3, e0) = op(e3, e1)), inference(cnf_transformation, [], [f3])).
fof(f3462, plain, (~ spl30_21 | spl30_53 | ~ spl30_123), inference(avatar_contradiction_clause, [], [f3461])).
fof(f3461, plain, ($false | (~ spl30_21 | spl30_53 | ~ spl30_123)), inference(subsumption_resolution, [], [f3456, f656])).
fof(f656, plain, (~ (e0 = op(e0, e2)) | spl30_53), inference(avatar_component_clause, [], [f655])).
fof(f3456, plain, ((e0 = op(e0, e2)) | (~ spl30_21 | ~ spl30_123)), inference(backward_demodulation, [], [f1175, f521])).
fof(f521, plain, ((e0 = op(e2, e2)) | ~ spl30_21), inference(avatar_component_clause, [], [f519])).
fof(f1175, plain, ((e0 = op(op(e2, e2), e2)) | ~ spl30_123), inference(avatar_component_clause, [], [f1174])).
fof(f1174, plain, (spl30_123 <=> (e0 = op(op(e2, e2), e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_123])])).
fof(f3460, plain, (~ spl30_17 | ~ spl30_21), inference(avatar_split_clause, [], [f3452, f519, f502])).
fof(f3452, plain, (~ (e0 = op(e2, e3)) | ~ spl30_21), inference(backward_demodulation, [], [f204, f521])).
fof(f204, plain, ~ (op(e2, e2) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f3459, plain, (~ spl30_5 | ~ spl30_21), inference(avatar_split_clause, [], [f3451, f519, f451])).
fof(f3451, plain, (~ (e0 = op(e3, e2)) | ~ spl30_21), inference(backward_demodulation, [], [f180, f521])).
fof(f3443, plain, (~ spl30_26 | ~ spl30_30), inference(avatar_split_clause, [], [f3436, f557, f540])).
fof(f540, plain, (spl30_26 <=> (e1 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_26])])).
fof(f3436, plain, (~ (e1 = op(e2, e1)) | ~ spl30_30), inference(backward_demodulation, [], [f199, f559])).
fof(f559, plain, ((e1 = op(e2, e0)) | ~ spl30_30), inference(avatar_component_clause, [], [f557])).
fof(f199, plain, ~ (op(e2, e0) = op(e2, e1)), inference(cnf_transformation, [], [f3])).
fof(f3420, plain, (~ spl30_26 | ~ spl30_42), inference(avatar_split_clause, [], [f3411, f608, f540])).
fof(f3411, plain, (~ (e1 = op(e2, e1)) | ~ spl30_42), inference(backward_demodulation, [], [f172, f610])).
fof(f3410, plain, (spl30_33 | ~ spl30_48 | ~ spl30_84), inference(avatar_split_clause, [], [f3408, f859, f633, f570])).
fof(f859, plain, (spl30_84 <=> (e0 = op(e1, op(e1, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl30_84])])).
fof(f3408, plain, ((e0 = op(e1, e3)) | (~ spl30_48 | ~ spl30_84)), inference(backward_demodulation, [], [f860, f635])).
fof(f635, plain, ((e3 = op(e1, e0)) | ~ spl30_48), inference(avatar_component_clause, [], [f633])).
fof(f860, plain, ((e0 = op(e1, op(e1, e0))) | ~ spl30_84), inference(avatar_component_clause, [], [f859])).
fof(f3409, plain, (~ spl30_44 | ~ spl30_48), inference(avatar_split_clause, [], [f3406, f633, f616])).
fof(f3406, plain, (~ (e3 = op(e1, e1)) | ~ spl30_48), inference(backward_demodulation, [], [f193, f635])).
fof(f193, plain, ~ (op(e1, e0) = op(e1, e1)), inference(cnf_transformation, [], [f3])).
fof(f3402, plain, (~ spl30_34 | ~ spl30_50), inference(avatar_split_clause, [], [f3397, f642, f574])).
fof(f3397, plain, (~ (e1 = op(e1, e3)) | ~ spl30_50), inference(backward_demodulation, [], [f181, f644])).
fof(f644, plain, ((e1 = op(e0, e3)) | ~ spl30_50), inference(avatar_component_clause, [], [f642])).
fof(f181, plain, ~ (op(e0, e3) = op(e1, e3)), inference(cnf_transformation, [], [f3])).
fof(f3396, plain, (~ spl30_51 | ~ spl30_56 | spl30_90), inference(avatar_split_clause, [], [f3394, f905, f667, f646])).
fof(f905, plain, (spl30_90 <=> (e2 = op(e0, op(e0, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl30_90])])).
fof(f3394, plain, (~ (e2 = op(e0, e3)) | (~ spl30_56 | spl30_90)), inference(backward_demodulation, [], [f907, f669])).
fof(f669, plain, ((e3 = op(e0, e2)) | ~ spl30_56), inference(avatar_component_clause, [], [f667])).
fof(f907, plain, (~ (e2 = op(e0, op(e0, e2))) | spl30_90), inference(avatar_component_clause, [], [f905])).
fof(f3395, plain, (~ spl30_24 | ~ spl30_56), inference(avatar_split_clause, [], [f3391, f667, f531])).
fof(f3391, plain, (~ (e3 = op(e2, e2)) | ~ spl30_56), inference(backward_demodulation, [], [f176, f669])).
fof(f3390, plain, (~ spl30_54 | ~ spl30_59 | spl30_91), inference(avatar_split_clause, [], [f3387, f910, f680, f659])).
fof(f659, plain, (spl30_54 <=> (e1 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_54])])).
fof(f680, plain, (spl30_59 <=> (e2 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_59])])).
fof(f910, plain, (spl30_91 <=> (e1 = op(e0, op(e0, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl30_91])])).
fof(f3387, plain, (~ (e1 = op(e0, e2)) | (~ spl30_59 | spl30_91)), inference(backward_demodulation, [], [f912, f682])).
fof(f682, plain, ((e2 = op(e0, e1)) | ~ spl30_59), inference(avatar_component_clause, [], [f680])).
fof(f912, plain, (~ (e1 = op(e0, op(e0, e1))) | spl30_91), inference(avatar_component_clause, [], [f910])).
fof(f3389, plain, (~ spl30_51 | ~ spl30_59), inference(avatar_split_clause, [], [f3385, f680, f646])).
fof(f3385, plain, (~ (e2 = op(e0, e3)) | ~ spl30_59), inference(backward_demodulation, [], [f191, f682])).
fof(f3388, plain, (~ spl30_11 | ~ spl30_59), inference(avatar_split_clause, [], [f3383, f680, f476])).
fof(f3383, plain, (~ (e2 = op(e3, e1)) | ~ spl30_59), inference(backward_demodulation, [], [f171, f682])).
fof(f3379, plain, (~ spl30_61 | ~ spl30_111), inference(avatar_contradiction_clause, [], [f3378])).
fof(f3378, plain, ($false | (~ spl30_61 | ~ spl30_111)), inference(subsumption_resolution, [], [f3377, f211])).
fof(f211, plain, ~ (e0 = e1), inference(cnf_transformation, [], [f4])).
fof(f4, plain, (~ (e2 = e3) & ~ (e1 = e3) & ~ (e1 = e2) & ~ (e0 = e3) & ~ (e0 = e2) & ~ (e0 = e1)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG159+1.p', ax4)).
fof(f3377, plain, ((e0 = e1) | (~ spl30_61 | ~ spl30_111)), inference(forward_demodulation, [], [f3372, f691])).
fof(f691, plain, ((e0 = op(e0, e0)) | ~ spl30_61), inference(avatar_component_clause, [], [f689])).
fof(f689, plain, (spl30_61 <=> (e0 = op(e0, e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_61])])).
fof(f3372, plain, ((op(e0, e0) = e1) | (~ spl30_61 | ~ spl30_111)), inference(backward_demodulation, [], [f1115, f691])).
fof(f1115, plain, ((e1 = op(op(e0, e0), e0)) | ~ spl30_111), inference(avatar_component_clause, [], [f1114])).
fof(f1114, plain, (spl30_111 <=> (e1 = op(op(e0, e0), e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_111])])).
fof(f3376, plain, (~ spl30_57 | ~ spl30_61), inference(avatar_split_clause, [], [f3365, f689, f672])).
fof(f3365, plain, (~ (e0 = op(e0, e1)) | ~ spl30_61), inference(backward_demodulation, [], [f187, f691])).
fof(f3375, plain, (~ spl30_45 | ~ spl30_61), inference(avatar_split_clause, [], [f3362, f689, f621])).
fof(f3362, plain, (~ (e0 = op(e1, e0)) | ~ spl30_61), inference(backward_demodulation, [], [f163, f691])).
fof(f3361, plain, (~ spl30_4 | spl30_66), inference(avatar_contradiction_clause, [], [f3360])).
fof(f3360, plain, ($false | (~ spl30_4 | spl30_66)), inference(subsumption_resolution, [], [f3359, f448])).
fof(f448, plain, ((e3 = op(e3, e3)) | ~ spl30_4), inference(avatar_component_clause, [], [f446])).
fof(f3359, plain, (~ (e3 = op(e3, e3)) | (~ spl30_4 | spl30_66)), inference(forward_demodulation, [], [f744, f448])).
fof(f744, plain, (~ (e3 = op(e3, op(e3, e3))) | spl30_66), inference(avatar_component_clause, [], [f742])).
fof(f742, plain, (spl30_66 <=> (e3 = op(e3, op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl30_66])])).
fof(f3329, plain, (~ spl30_3 | ~ spl30_11), inference(avatar_split_clause, [], [f3324, f476, f442])).
fof(f442, plain, (spl30_3 <=> (e2 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_3])])).
fof(f3324, plain, (~ (e2 = op(e3, e3)) | ~ spl30_11), inference(backward_demodulation, [], [f209, f478])).
fof(f478, plain, ((e2 = op(e3, e1)) | ~ spl30_11), inference(avatar_component_clause, [], [f476])).
fof(f209, plain, ~ (op(e3, e1) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f3322, plain, (~ spl30_10 | ~ spl30_14), inference(avatar_split_clause, [], [f3316, f489, f472])).
fof(f472, plain, (spl30_10 <=> (e1 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_10])])).
fof(f489, plain, (spl30_14 <=> (e1 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_14])])).
fof(f3316, plain, (~ (e1 = op(e3, e1)) | ~ spl30_14), inference(backward_demodulation, [], [f205, f491])).
fof(f491, plain, ((e1 = op(e3, e0)) | ~ spl30_14), inference(avatar_component_clause, [], [f489])).
fof(f3311, plain, (~ spl30_14 | ~ spl30_24 | ~ spl30_123 | spl30_125), inference(avatar_split_clause, [], [f3305, f1183, f1174, f531, f489])).
fof(f1183, plain, (spl30_125 <=> (e1 = op(op(e2, e2), op(op(e2, e2), e2)))), introduced(avatar_definition, [new_symbols(naming, [spl30_125])])).
fof(f3305, plain, (~ (e1 = op(e3, e0)) | (~ spl30_24 | ~ spl30_123 | spl30_125)), inference(backward_demodulation, [], [f3215, f533])).
fof(f3215, plain, (~ (e1 = op(op(e2, e2), e0)) | (~ spl30_123 | spl30_125)), inference(backward_demodulation, [], [f1185, f1175])).
fof(f1185, plain, (~ (e1 = op(op(e2, e2), op(op(e2, e2), e2))) | spl30_125), inference(avatar_component_clause, [], [f1183])).
fof(f3309, plain, (spl30_19 | ~ spl30_24 | ~ spl30_74), inference(avatar_contradiction_clause, [], [f3308])).
fof(f3308, plain, ($false | (spl30_19 | ~ spl30_24 | ~ spl30_74)), inference(subsumption_resolution, [], [f3303, f511])).
fof(f511, plain, (~ (e2 = op(e2, e3)) | spl30_19), inference(avatar_component_clause, [], [f510])).
fof(f3303, plain, ((e2 = op(e2, e3)) | (~ spl30_24 | ~ spl30_74)), inference(backward_demodulation, [], [f794, f533])).
fof(f794, plain, ((e2 = op(e2, op(e2, e2))) | ~ spl30_74), inference(avatar_component_clause, [], [f793])).
fof(f3307, plain, (~ spl30_8 | ~ spl30_24), inference(avatar_split_clause, [], [f3301, f531, f463])).
fof(f463, plain, (spl30_8 <=> (e3 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_8])])).
fof(f3301, plain, (~ (e3 = op(e3, e2)) | ~ spl30_24), inference(backward_demodulation, [], [f180, f533])).
fof(f3300, plain, (~ spl30_18 | ~ spl30_26), inference(avatar_split_clause, [], [f3295, f540, f506])).
fof(f506, plain, (spl30_18 <=> (e1 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_18])])).
fof(f3295, plain, (~ (e1 = op(e2, e3)) | ~ spl30_26), inference(backward_demodulation, [], [f203, f542])).
fof(f542, plain, ((e1 = op(e2, e1)) | ~ spl30_26), inference(avatar_component_clause, [], [f540])).
fof(f3299, plain, (~ spl30_10 | ~ spl30_26), inference(avatar_split_clause, [], [f3293, f540, f472])).
fof(f3293, plain, (~ (e1 = op(e3, e1)) | ~ spl30_26), inference(backward_demodulation, [], [f174, f542])).
fof(f174, plain, ~ (op(e2, e1) = op(e3, e1)), inference(cnf_transformation, [], [f3])).
fof(f3292, plain, (~ spl30_23 | ~ spl30_31), inference(avatar_split_clause, [], [f3288, f561, f527])).
fof(f527, plain, (spl30_23 <=> (e2 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_23])])).
fof(f561, plain, (spl30_31 <=> (e2 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_31])])).
fof(f3288, plain, (~ (e2 = op(e2, e2)) | ~ spl30_31), inference(backward_demodulation, [], [f200, f563])).
fof(f563, plain, ((e2 = op(e2, e0)) | ~ spl30_31), inference(avatar_component_clause, [], [f561])).
fof(f200, plain, ~ (op(e2, e0) = op(e2, e2)), inference(cnf_transformation, [], [f3])).
fof(f3285, plain, (~ spl30_18 | ~ spl30_34), inference(avatar_split_clause, [], [f3280, f574, f506])).
fof(f3280, plain, (~ (e1 = op(e2, e3)) | ~ spl30_34), inference(backward_demodulation, [], [f184, f576])).
fof(f576, plain, ((e1 = op(e1, e3)) | ~ spl30_34), inference(avatar_component_clause, [], [f574])).
fof(f3279, plain, (~ spl30_23 | ~ spl30_39), inference(avatar_split_clause, [], [f3274, f595, f527])).
fof(f3274, plain, (~ (e2 = op(e2, e2)) | ~ spl30_39), inference(backward_demodulation, [], [f178, f597])).
fof(f597, plain, ((e2 = op(e1, e2)) | ~ spl30_39), inference(avatar_component_clause, [], [f595])).
fof(f3273, plain, (~ spl30_34 | ~ spl30_44 | spl30_83), inference(avatar_split_clause, [], [f3268, f854, f616, f574])).
fof(f3268, plain, (~ (e1 = op(e1, e3)) | (~ spl30_44 | spl30_83)), inference(backward_demodulation, [], [f856, f618])).
fof(f856, plain, (~ (e1 = op(e1, op(e1, e1))) | spl30_83), inference(avatar_component_clause, [], [f854])).
fof(f3272, plain, (~ spl30_36 | ~ spl30_44), inference(avatar_split_clause, [], [f3267, f616, f582])).
fof(f582, plain, (spl30_36 <=> (e3 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_36])])).
fof(f3267, plain, (~ (e3 = op(e1, e3)) | ~ spl30_44), inference(backward_demodulation, [], [f197, f618])).
fof(f3263, plain, (~ spl30_37 | ~ spl30_45), inference(avatar_split_clause, [], [f3256, f621, f587])).
fof(f587, plain, (spl30_37 <=> (e0 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_37])])).
fof(f3256, plain, (~ (e0 = op(e1, e2)) | ~ spl30_45), inference(backward_demodulation, [], [f194, f623])).
fof(f3262, plain, (~ spl30_13 | ~ spl30_45), inference(avatar_split_clause, [], [f3254, f621, f485])).
fof(f485, plain, (spl30_13 <=> (e0 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_13])])).
fof(f3254, plain, (~ (e0 = op(e3, e0)) | ~ spl30_45), inference(backward_demodulation, [], [f167, f623])).
fof(f3251, plain, (~ spl30_3 | ~ spl30_51), inference(avatar_split_clause, [], [f3248, f646, f442])).
fof(f3248, plain, (~ (e2 = op(e3, e3)) | ~ spl30_51), inference(backward_demodulation, [], [f183, f648])).
fof(f3245, plain, (~ spl30_49 | ~ spl30_57), inference(avatar_split_clause, [], [f3239, f672, f638])).
fof(f638, plain, (spl30_49 <=> (e0 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_49])])).
fof(f3239, plain, (~ (e0 = op(e0, e3)) | ~ spl30_57), inference(backward_demodulation, [], [f191, f674])).
fof(f3244, plain, (~ spl30_25 | ~ spl30_57), inference(avatar_split_clause, [], [f3237, f672, f536])).
fof(f536, plain, (spl30_25 <=> (e0 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_25])])).
fof(f3237, plain, (~ (e0 = op(e2, e1)) | ~ spl30_57), inference(backward_demodulation, [], [f170, f674])).
fof(f170, plain, ~ (op(e0, e1) = op(e2, e1)), inference(cnf_transformation, [], [f3])).
fof(f3233, plain, (~ spl30_49 | ~ spl30_64 | spl30_92), inference(avatar_split_clause, [], [f3227, f915, f701, f638])).
fof(f915, plain, (spl30_92 <=> (e0 = op(e0, op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl30_92])])).
fof(f3227, plain, (~ (e0 = op(e0, e3)) | (~ spl30_64 | spl30_92)), inference(backward_demodulation, [], [f917, f703])).
fof(f917, plain, (~ (e0 = op(e0, op(e0, e0))) | spl30_92), inference(avatar_component_clause, [], [f915])).
fof(f3232, plain, (~ spl30_60 | ~ spl30_64), inference(avatar_split_clause, [], [f3225, f701, f684])).
fof(f684, plain, (spl30_60 <=> (e3 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_60])])).
fof(f3225, plain, (~ (e3 = op(e0, e1)) | ~ spl30_64), inference(backward_demodulation, [], [f187, f703])).
fof(f3231, plain, (~ spl30_32 | ~ spl30_64), inference(avatar_split_clause, [], [f3223, f701, f565])).
fof(f565, plain, (spl30_32 <=> (e3 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_32])])).
fof(f3223, plain, (~ (e3 = op(e2, e0)) | ~ spl30_64), inference(backward_demodulation, [], [f164, f703])).
fof(f3208, plain, (~ spl30_5 | ~ spl30_13), inference(avatar_split_clause, [], [f3203, f485, f451])).
fof(f3203, plain, (~ (e0 = op(e3, e2)) | ~ spl30_13), inference(backward_demodulation, [], [f206, f487])).
fof(f487, plain, ((e0 = op(e3, e0)) | ~ spl30_13), inference(avatar_component_clause, [], [f485])).
fof(f206, plain, ~ (op(e3, e0) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f3192, plain, (~ spl30_17 | ~ spl30_25), inference(avatar_split_clause, [], [f3185, f536, f502])).
fof(f3185, plain, (~ (e0 = op(e2, e3)) | ~ spl30_25), inference(backward_demodulation, [], [f203, f538])).
fof(f538, plain, ((e0 = op(e2, e1)) | ~ spl30_25), inference(avatar_component_clause, [], [f536])).
fof(f3183, plain, (~ spl30_20 | ~ spl30_32), inference(avatar_split_clause, [], [f3182, f565, f514])).
fof(f514, plain, (spl30_20 <=> (e3 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_20])])).
fof(f3182, plain, (~ (e3 = op(e2, e3)) | ~ spl30_32), inference(backward_demodulation, [], [f201, f567])).
fof(f567, plain, ((e3 = op(e2, e0)) | ~ spl30_32), inference(avatar_component_clause, [], [f565])).
fof(f3179, plain, (~ spl30_20 | ~ spl30_36), inference(avatar_split_clause, [], [f3177, f582, f514])).
fof(f3177, plain, (~ (e3 = op(e2, e3)) | ~ spl30_36), inference(backward_demodulation, [], [f184, f584])).
fof(f584, plain, ((e3 = op(e1, e3)) | ~ spl30_36), inference(avatar_component_clause, [], [f582])).
fof(f3175, plain, (~ spl30_5 | ~ spl30_37), inference(avatar_split_clause, [], [f3169, f587, f451])).
fof(f3169, plain, (~ (e0 = op(e3, e2)) | ~ spl30_37), inference(backward_demodulation, [], [f179, f589])).
fof(f589, plain, ((e0 = op(e1, e2)) | ~ spl30_37), inference(avatar_component_clause, [], [f587])).
fof(f179, plain, ~ (op(e1, e2) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f3168, plain, (~ spl30_34 | ~ spl30_46), inference(avatar_split_clause, [], [f3163, f625, f574])).
fof(f625, plain, (spl30_46 <=> (e1 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_46])])).
fof(f3163, plain, (~ (e1 = op(e1, e3)) | ~ spl30_46), inference(backward_demodulation, [], [f195, f627])).
fof(f627, plain, ((e1 = op(e1, e0)) | ~ spl30_46), inference(avatar_component_clause, [], [f625])).
fof(f3167, plain, (~ spl30_14 | ~ spl30_46), inference(avatar_split_clause, [], [f3161, f625, f489])).
fof(f3161, plain, (~ (e1 = op(e3, e0)) | ~ spl30_46), inference(backward_demodulation, [], [f167, f627])).
fof(f3158, plain, (~ spl30_17 | ~ spl30_49), inference(avatar_split_clause, [], [f3153, f638, f502])).
fof(f3153, plain, (~ (e0 = op(e2, e3)) | ~ spl30_49), inference(backward_demodulation, [], [f182, f640])).
fof(f640, plain, ((e0 = op(e0, e3)) | ~ spl30_49), inference(avatar_component_clause, [], [f638])).
fof(f3151, plain, (~ spl30_12 | ~ spl30_60), inference(avatar_split_clause, [], [f3148, f684, f480])).
fof(f480, plain, (spl30_12 <=> (e3 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_12])])).
fof(f3148, plain, (~ (e3 = op(e3, e1)) | ~ spl30_60), inference(backward_demodulation, [], [f171, f686])).
fof(f686, plain, ((e3 = op(e0, e1)) | ~ spl30_60), inference(avatar_component_clause, [], [f684])).
fof(f3146, plain, (~ spl30_18 | ~ spl30_63 | spl30_115 | ~ spl30_116), inference(avatar_split_clause, [], [f3140, f1138, f1133, f697, f506])).
fof(f697, plain, (spl30_63 <=> (op(e0, e0) = e2)), introduced(avatar_definition, [new_symbols(naming, [spl30_63])])).
fof(f1138, plain, (spl30_116 <=> (e3 = op(op(e0, e0), e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_116])])).
fof(f3140, plain, (~ (e1 = op(e2, e3)) | (~ spl30_63 | spl30_115 | ~ spl30_116)), inference(backward_demodulation, [], [f3128, f699])).
fof(f699, plain, ((op(e0, e0) = e2) | ~ spl30_63), inference(avatar_component_clause, [], [f697])).
fof(f3128, plain, (~ (e1 = op(op(e0, e0), e3)) | (spl30_115 | ~ spl30_116)), inference(forward_demodulation, [], [f1135, f1139])).
fof(f1139, plain, ((e3 = op(op(e0, e0), e0)) | ~ spl30_116), inference(avatar_component_clause, [], [f1138])).
fof(f3141, plain, (~ spl30_31 | ~ spl30_63), inference(avatar_split_clause, [], [f3130, f697, f561])).
fof(f3130, plain, (~ (e2 = op(e2, e0)) | ~ spl30_63), inference(backward_demodulation, [], [f164, f699])).
fof(f3117, plain, (~ spl30_27 | ~ spl30_43), inference(avatar_split_clause, [], [f3116, f612, f544])).
fof(f612, plain, (spl30_43 <=> (e2 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_43])])).
fof(f3116, plain, (~ (e2 = op(e2, e1)) | ~ spl30_43), inference(forward_demodulation, [], [f172, f614])).
fof(f614, plain, ((e2 = op(e1, e1)) | ~ spl30_43), inference(avatar_component_clause, [], [f612])).
fof(f3115, plain, (~ spl30_25 | ~ spl30_43 | spl30_117), inference(avatar_split_clause, [], [f3076, f1144, f612, f536])).
fof(f3076, plain, (~ (e0 = op(e2, e1)) | (~ spl30_43 | spl30_117)), inference(backward_demodulation, [], [f1146, f614])).
fof(f3112, plain, (~ spl30_6 | ~ spl30_54), inference(avatar_split_clause, [], [f3111, f659, f455])).
fof(f3111, plain, (~ (e1 = op(e3, e2)) | ~ spl30_54), inference(forward_demodulation, [], [f177, f661])).
fof(f661, plain, ((e1 = op(e0, e2)) | ~ spl30_54), inference(avatar_component_clause, [], [f659])).
fof(f177, plain, ~ (op(e0, e2) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f3110, plain, (~ spl30_2 | ~ spl30_12 | spl30_68), inference(avatar_split_clause, [], [f3109, f752, f480, f438])).
fof(f752, plain, (spl30_68 <=> (e1 = op(e3, op(e3, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl30_68])])).
fof(f3109, plain, (~ (e1 = op(e3, e3)) | (~ spl30_12 | spl30_68)), inference(backward_demodulation, [], [f754, f482])).
fof(f482, plain, ((e3 = op(e3, e1)) | ~ spl30_12), inference(avatar_component_clause, [], [f480])).
fof(f754, plain, (~ (e1 = op(e3, op(e3, e1))) | spl30_68), inference(avatar_component_clause, [], [f752])).
fof(f3108, plain, (~ spl30_9 | ~ spl30_14 | spl30_69), inference(avatar_split_clause, [], [f3107, f757, f489, f468])).
fof(f3107, plain, (~ (e0 = op(e3, e1)) | (~ spl30_14 | spl30_69)), inference(backward_demodulation, [], [f759, f491])).
fof(f3103, plain, (~ spl30_17 | ~ spl30_20), inference(avatar_contradiction_clause, [], [f3102])).
fof(f3102, plain, ($false | (~ spl30_17 | ~ spl30_20)), inference(subsumption_resolution, [], [f3101, f213])).
fof(f213, plain, ~ (e0 = e3), inference(cnf_transformation, [], [f4])).
fof(f3101, plain, ((e0 = e3) | (~ spl30_17 | ~ spl30_20)), inference(backward_demodulation, [], [f516, f504])).
fof(f504, plain, ((e0 = op(e2, e3)) | ~ spl30_17), inference(avatar_component_clause, [], [f502])).
fof(f516, plain, ((e3 = op(e2, e3)) | ~ spl30_20), inference(avatar_component_clause, [], [f514])).
fof(f3081, plain, (~ spl30_38 | ~ spl30_43 | spl30_83), inference(avatar_split_clause, [], [f3075, f854, f612, f591])).
fof(f3075, plain, (~ (e1 = op(e1, e2)) | (~ spl30_43 | spl30_83)), inference(backward_demodulation, [], [f856, f614])).
fof(f3069, plain, (~ spl30_29 | ~ spl30_45), inference(avatar_split_clause, [], [f3060, f621, f553])).
fof(f3060, plain, (~ (e0 = op(e2, e0)) | ~ spl30_45), inference(backward_demodulation, [], [f166, f623])).
fof(f166, plain, ~ (op(e1, e0) = op(e2, e0)), inference(cnf_transformation, [], [f3])).
fof(f3054, plain, (~ spl30_38 | ~ spl30_54), inference(avatar_split_clause, [], [f3049, f659, f591])).
fof(f3049, plain, (~ (e1 = op(e1, e2)) | ~ spl30_54), inference(backward_demodulation, [], [f175, f661])).
fof(f3034, plain, (~ spl30_16 | ~ spl30_64), inference(avatar_split_clause, [], [f3026, f701, f497])).
fof(f3026, plain, (~ (e3 = op(e3, e0)) | ~ spl30_64), inference(backward_demodulation, [], [f165, f703])).
fof(f165, plain, ~ (op(e0, e0) = op(e3, e0)), inference(cnf_transformation, [], [f3])).
fof(f3009, plain, (~ spl30_2 | ~ spl30_3), inference(avatar_contradiction_clause, [], [f3008])).
fof(f3008, plain, ($false | (~ spl30_2 | ~ spl30_3)), inference(subsumption_resolution, [], [f3007, f214])).
fof(f214, plain, ~ (e1 = e2), inference(cnf_transformation, [], [f4])).
fof(f3007, plain, ((e1 = e2) | (~ spl30_2 | ~ spl30_3)), inference(backward_demodulation, [], [f444, f440])).
fof(f444, plain, ((e2 = op(e3, e3)) | ~ spl30_3), inference(avatar_component_clause, [], [f442])).
fof(f2999, plain, (~ spl30_12 | ~ spl30_16), inference(avatar_split_clause, [], [f2997, f497, f480])).
fof(f2997, plain, (~ (e3 = op(e3, e1)) | ~ spl30_16), inference(backward_demodulation, [], [f205, f499])).
fof(f2990, plain, (~ spl30_12 | ~ spl30_44), inference(avatar_split_clause, [], [f2983, f616, f480])).
fof(f2983, plain, (~ (e3 = op(e3, e1)) | ~ spl30_44), inference(backward_demodulation, [], [f173, f618])).
fof(f173, plain, ~ (op(e1, e1) = op(e3, e1)), inference(cnf_transformation, [], [f3])).
fof(f2979, plain, (~ spl30_35 | ~ spl30_47), inference(avatar_split_clause, [], [f2975, f629, f578])).
fof(f578, plain, (spl30_35 <=> (e2 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_35])])).
fof(f2975, plain, (~ (e2 = op(e1, e3)) | ~ spl30_47), inference(backward_demodulation, [], [f195, f631])).
fof(f2964, plain, (~ spl30_35 | ~ spl30_3), inference(avatar_split_clause, [], [f2963, f442, f578])).
fof(f2963, plain, (~ (e2 = op(e1, e3)) | ~ spl30_3), inference(forward_demodulation, [], [f185, f444])).
fof(f185, plain, ~ (op(e1, e3) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f2962, plain, (~ spl30_33 | ~ spl30_49), inference(avatar_split_clause, [], [f2865, f638, f570])).
fof(f2865, plain, (~ (e0 = op(e1, e3)) | ~ spl30_49), inference(backward_demodulation, [], [f181, f640])).
fof(f2933, plain, (~ spl30_1 | ~ spl30_5), inference(avatar_split_clause, [], [f2928, f451, f434])).
fof(f434, plain, (spl30_1 <=> (e0 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_1])])).
fof(f2928, plain, (~ (e0 = op(e3, e3)) | ~ spl30_5), inference(backward_demodulation, [], [f210, f453])).
fof(f453, plain, ((e0 = op(e3, e2)) | ~ spl30_5), inference(avatar_component_clause, [], [f451])).
fof(f210, plain, ~ (op(e3, e2) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f2917, plain, (~ spl30_18 | ~ spl30_20), inference(avatar_contradiction_clause, [], [f2916])).
fof(f2916, plain, ($false | (~ spl30_18 | ~ spl30_20)), inference(subsumption_resolution, [], [f2915, f215])).
fof(f215, plain, ~ (e1 = e3), inference(cnf_transformation, [], [f4])).
fof(f2915, plain, ((e1 = e3) | (~ spl30_18 | ~ spl30_20)), inference(backward_demodulation, [], [f516, f508])).
fof(f508, plain, ((e1 = op(e2, e3)) | ~ spl30_18), inference(avatar_component_clause, [], [f506])).
fof(f2913, plain, (~ spl30_29 | ~ spl30_31), inference(avatar_contradiction_clause, [], [f2912])).
fof(f2912, plain, ($false | (~ spl30_29 | ~ spl30_31)), inference(subsumption_resolution, [], [f2911, f212])).
fof(f212, plain, ~ (e0 = e2), inference(cnf_transformation, [], [f4])).
fof(f2911, plain, ((e0 = e2) | (~ spl30_29 | ~ spl30_31)), inference(forward_demodulation, [], [f563, f555])).
fof(f2904, plain, (~ spl30_34 | ~ spl30_38), inference(avatar_split_clause, [], [f2900, f591, f574])).
fof(f2900, plain, (~ (e1 = op(e1, e3)) | ~ spl30_38), inference(backward_demodulation, [], [f198, f593])).
fof(f593, plain, ((e1 = op(e1, e2)) | ~ spl30_38), inference(avatar_component_clause, [], [f591])).
fof(f198, plain, ~ (op(e1, e2) = op(e1, e3)), inference(cnf_transformation, [], [f3])).
fof(f2897, plain, (~ spl30_41 | ~ spl30_59 | spl30_120), inference(avatar_contradiction_clause, [], [f2896])).
fof(f2896, plain, ($false | (~ spl30_41 | ~ spl30_59 | spl30_120)), inference(subsumption_resolution, [], [f2891, f682])).
fof(f2891, plain, (~ (e2 = op(e0, e1)) | (~ spl30_41 | spl30_120)), inference(backward_demodulation, [], [f1160, f606])).
fof(f606, plain, ((e0 = op(e1, e1)) | ~ spl30_41), inference(avatar_component_clause, [], [f604])).
fof(f604, plain, (spl30_41 <=> (e0 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_41])])).
fof(f1160, plain, (~ (e2 = op(op(e1, e1), e1)) | spl30_120), inference(avatar_component_clause, [], [f1158])).
fof(f1158, plain, (spl30_120 <=> (e2 = op(op(e1, e1), e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_120])])).
fof(f2893, plain, (~ spl30_37 | ~ spl30_41), inference(avatar_split_clause, [], [f2883, f604, f587])).
fof(f2883, plain, (~ (e0 = op(e1, e2)) | ~ spl30_41), inference(backward_demodulation, [], [f196, f606])).
fof(f196, plain, ~ (op(e1, e1) = op(e1, e2)), inference(cnf_transformation, [], [f3])).
fof(f2880, plain, (~ spl30_40 | ~ spl30_48), inference(avatar_split_clause, [], [f2875, f633, f599])).
fof(f2875, plain, (~ (e3 = op(e1, e2)) | ~ spl30_48), inference(backward_demodulation, [], [f194, f635])).
fof(f2879, plain, (~ spl30_16 | ~ spl30_48), inference(avatar_split_clause, [], [f2873, f633, f497])).
fof(f2873, plain, (~ (e3 = op(e3, e0)) | ~ spl30_48), inference(backward_demodulation, [], [f167, f635])).
fof(f2871, plain, (~ spl30_1 | ~ spl30_49), inference(avatar_split_clause, [], [f2866, f638, f434])).
fof(f2866, plain, (~ (e0 = op(e3, e3)) | ~ spl30_49), inference(backward_demodulation, [], [f183, f640])).
fof(f2856, plain, (~ spl30_43 | ~ spl30_59), inference(avatar_split_clause, [], [f2850, f680, f612])).
fof(f2850, plain, (~ (e2 = op(e1, e1)) | ~ spl30_59), inference(backward_demodulation, [], [f169, f682])).
fof(f2849, plain, (~ spl30_48 | ~ spl30_62 | spl30_116), inference(avatar_split_clause, [], [f2843, f1138, f693, f633])).
fof(f2843, plain, (~ (e3 = op(e1, e0)) | (~ spl30_62 | spl30_116)), inference(backward_demodulation, [], [f1140, f695])).
fof(f1140, plain, (~ (e3 = op(op(e0, e0), e0)) | spl30_116), inference(avatar_component_clause, [], [f1138])).
fof(f2847, plain, (~ spl30_57 | ~ spl30_62 | spl30_92), inference(avatar_split_clause, [], [f2840, f915, f693, f672])).
fof(f2840, plain, (~ (e0 = op(e0, e1)) | (~ spl30_62 | spl30_92)), inference(backward_demodulation, [], [f917, f695])).
fof(f2846, plain, (~ spl30_54 | ~ spl30_62), inference(avatar_split_clause, [], [f2836, f693, f659])).
fof(f2836, plain, (~ (e1 = op(e0, e2)) | ~ spl30_62), inference(backward_demodulation, [], [f188, f695])).
fof(f2845, plain, (~ spl30_14 | ~ spl30_62), inference(avatar_split_clause, [], [f2834, f693, f489])).
fof(f2834, plain, (~ (e1 = op(e3, e0)) | ~ spl30_62), inference(backward_demodulation, [], [f165, f695])).
fof(f2844, plain, (~ spl30_46 | ~ spl30_62), inference(avatar_split_clause, [], [f2833, f693, f625])).
fof(f2833, plain, (~ (e1 = op(e1, e0)) | ~ spl30_62), inference(backward_demodulation, [], [f163, f695])).
fof(f2818, plain, (~ spl30_7 | ~ spl30_23), inference(avatar_split_clause, [], [f2817, f527, f459])).
fof(f459, plain, (spl30_7 <=> (e2 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_7])])).
fof(f2817, plain, (~ (e2 = op(e3, e2)) | ~ spl30_23), inference(forward_demodulation, [], [f180, f529])).
fof(f529, plain, ((e2 = op(e2, e2)) | ~ spl30_23), inference(avatar_component_clause, [], [f527])).
fof(f2815, plain, (~ spl30_1 | ~ spl30_54 | spl30_131 | ~ spl30_134), inference(avatar_contradiction_clause, [], [f2814])).
fof(f2814, plain, ($false | (~ spl30_1 | ~ spl30_54 | spl30_131 | ~ spl30_134)), inference(subsumption_resolution, [], [f2811, f661])).
fof(f2811, plain, (~ (e1 = op(e0, e2)) | (~ spl30_1 | spl30_131 | ~ spl30_134)), inference(backward_demodulation, [], [f2688, f436])).
fof(f436, plain, ((e0 = op(e3, e3)) | ~ spl30_1), inference(avatar_component_clause, [], [f434])).
fof(f2688, plain, (~ (e1 = op(op(e3, e3), e2)) | (spl30_131 | ~ spl30_134)), inference(backward_demodulation, [], [f1215, f1229])).
fof(f1229, plain, ((e2 = op(op(e3, e3), e3)) | ~ spl30_134), inference(avatar_component_clause, [], [f1228])).
fof(f1228, plain, (spl30_134 <=> (e2 = op(op(e3, e3), e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_134])])).
fof(f1215, plain, (~ (e1 = op(op(e3, e3), op(op(e3, e3), e3))) | spl30_131), inference(avatar_component_clause, [], [f1213])).
fof(f1213, plain, (spl30_131 <=> (e1 = op(op(e3, e3), op(op(e3, e3), e3)))), introduced(avatar_definition, [new_symbols(naming, [spl30_131])])).
fof(f2804, plain, (~ spl30_2 | ~ spl30_10), inference(avatar_split_clause, [], [f2800, f472, f438])).
fof(f2800, plain, (~ (e1 = op(e3, e3)) | ~ spl30_10), inference(backward_demodulation, [], [f209, f474])).
fof(f474, plain, ((e1 = op(e3, e1)) | ~ spl30_10), inference(avatar_component_clause, [], [f472])).
fof(f2787, plain, (~ spl30_21 | ~ spl30_23 | spl30_123), inference(avatar_split_clause, [], [f2780, f1174, f527, f519])).
fof(f2780, plain, (~ (e0 = op(e2, e2)) | (~ spl30_23 | spl30_123)), inference(backward_demodulation, [], [f1176, f529])).
fof(f1176, plain, (~ (e0 = op(op(e2, e2), e2)) | spl30_123), inference(avatar_component_clause, [], [f1174])).
fof(f2786, plain, (~ spl30_19 | ~ spl30_23), inference(avatar_split_clause, [], [f2777, f527, f510])).
fof(f2777, plain, (~ (e2 = op(e2, e3)) | ~ spl30_23), inference(backward_demodulation, [], [f204, f529])).
fof(f2776, plain, (~ spl30_22 | ~ spl30_26), inference(avatar_split_clause, [], [f2769, f540, f523])).
fof(f2769, plain, (~ (e1 = op(e2, e2)) | ~ spl30_26), inference(backward_demodulation, [], [f202, f542])).
fof(f2767, plain, (~ spl30_21 | ~ spl30_29), inference(avatar_split_clause, [], [f2761, f553, f519])).
fof(f2761, plain, (~ (e0 = op(e2, e2)) | ~ spl30_29), inference(backward_demodulation, [], [f200, f555])).
fof(f2758, plain, (~ spl30_2 | ~ spl30_34), inference(avatar_split_clause, [], [f2754, f574, f438])).
fof(f2754, plain, (~ (e1 = op(e3, e3)) | ~ spl30_34), inference(backward_demodulation, [], [f185, f576])).
fof(f2752, plain, (~ spl30_37 | ~ spl30_40), inference(avatar_contradiction_clause, [], [f2751])).
fof(f2751, plain, ($false | (~ spl30_37 | ~ spl30_40)), inference(subsumption_resolution, [], [f2750, f213])).
fof(f2750, plain, ((e0 = e3) | (~ spl30_37 | ~ spl30_40)), inference(backward_demodulation, [], [f601, f589])).
fof(f2747, plain, (~ spl30_36 | ~ spl30_40), inference(avatar_split_clause, [], [f2745, f599, f582])).
fof(f2745, plain, (~ (e3 = op(e1, e3)) | ~ spl30_40), inference(backward_demodulation, [], [f198, f601])).
fof(f2743, plain, (~ spl30_28 | ~ spl30_43 | spl30_122), inference(avatar_split_clause, [], [f2739, f1168, f612, f548])).
fof(f1168, plain, (spl30_122 <=> (e3 = op(op(e1, e1), e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_122])])).
fof(f2739, plain, (~ (e3 = op(e2, e1)) | (~ spl30_43 | spl30_122)), inference(backward_demodulation, [], [f1170, f614])).
fof(f1170, plain, (~ (e3 = op(op(e1, e1), e1)) | spl30_122), inference(avatar_component_clause, [], [f1168])).
fof(f2740, plain, (~ spl30_39 | ~ spl30_43), inference(avatar_split_clause, [], [f2732, f612, f595])).
fof(f2732, plain, (~ (e2 = op(e1, e2)) | ~ spl30_43), inference(backward_demodulation, [], [f196, f614])).
fof(f2728, plain, (~ spl30_42 | ~ spl30_46), inference(avatar_split_clause, [], [f2722, f625, f608])).
fof(f2722, plain, (~ (e1 = op(e1, e1)) | ~ spl30_46), inference(backward_demodulation, [], [f193, f627])).
fof(f2712, plain, (~ spl30_22 | ~ spl30_54), inference(avatar_split_clause, [], [f2707, f659, f523])).
fof(f2707, plain, (~ (e1 = op(e2, e2)) | ~ spl30_54), inference(backward_demodulation, [], [f176, f661])).
fof(f2705, plain, (~ spl30_57 | ~ spl30_60), inference(avatar_contradiction_clause, [], [f2704])).
fof(f2704, plain, ($false | (~ spl30_57 | ~ spl30_60)), inference(subsumption_resolution, [], [f2703, f213])).
fof(f2703, plain, ((e0 = e3) | (~ spl30_57 | ~ spl30_60)), inference(backward_demodulation, [], [f686, f674])).
fof(f2700, plain, (~ spl30_56 | ~ spl30_60), inference(avatar_split_clause, [], [f2696, f684, f667])).
fof(f2696, plain, (~ (e3 = op(e0, e2)) | ~ spl30_60), inference(backward_demodulation, [], [f190, f686])).
fof(f2699, plain, (~ spl30_28 | ~ spl30_60), inference(avatar_split_clause, [], [f2694, f684, f548])).
fof(f2694, plain, (~ (e3 = op(e2, e1)) | ~ spl30_60), inference(backward_demodulation, [], [f170, f686])).
fof(f2679, plain, (~ spl30_59 | ~ spl30_63), inference(avatar_split_clause, [], [f2678, f697, f680])).
fof(f2678, plain, (~ (e2 = op(e0, e1)) | ~ spl30_63), inference(forward_demodulation, [], [f187, f699])).
fof(f2674, plain, (~ spl30_32 | ~ spl30_63 | spl30_116), inference(avatar_split_clause, [], [f2673, f1138, f697, f565])).
fof(f2673, plain, (~ (e3 = op(e2, e0)) | (~ spl30_63 | spl30_116)), inference(forward_demodulation, [], [f1140, f699])).
fof(f2670, plain, (~ spl30_39 | ~ spl30_7), inference(avatar_split_clause, [], [f2669, f459, f595])).
fof(f2669, plain, (~ (e2 = op(e1, e2)) | ~ spl30_7), inference(forward_demodulation, [], [f179, f461])).
fof(f461, plain, ((e2 = op(e3, e2)) | ~ spl30_7), inference(avatar_component_clause, [], [f459])).
fof(f2668, plain, (~ spl30_32 | ~ spl30_16), inference(avatar_split_clause, [], [f2667, f497, f565])).
fof(f2667, plain, (~ (e3 = op(e2, e0)) | ~ spl30_16), inference(forward_demodulation, [], [f168, f499])).
fof(f168, plain, ~ (op(e2, e0) = op(e3, e0)), inference(cnf_transformation, [], [f3])).
fof(f2663, plain, (spl30_1 | ~ spl30_16 | ~ spl30_69), inference(avatar_split_clause, [], [f2641, f757, f497, f434])).
fof(f2641, plain, ((e0 = op(e3, e3)) | (~ spl30_16 | ~ spl30_69)), inference(backward_demodulation, [], [f758, f499])).
fof(f758, plain, ((e0 = op(e3, op(e3, e0))) | ~ spl30_69), inference(avatar_component_clause, [], [f757])).
fof(f2655, plain, (~ spl30_7 | spl30_67), inference(avatar_contradiction_clause, [], [f2654])).
fof(f2654, plain, ($false | (~ spl30_7 | spl30_67)), inference(subsumption_resolution, [], [f2652, f461])).
fof(f2652, plain, (~ (e2 = op(e3, e2)) | (~ spl30_7 | spl30_67)), inference(backward_demodulation, [], [f749, f461])).
fof(f749, plain, (~ (e2 = op(e3, op(e3, e2))) | spl30_67), inference(avatar_component_clause, [], [f747])).
fof(f747, plain, (spl30_67 <=> (e2 = op(e3, op(e3, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl30_67])])).
fof(f2653, plain, (~ spl30_3 | ~ spl30_7), inference(avatar_split_clause, [], [f2650, f459, f442])).
fof(f2650, plain, (~ (e2 = op(e3, e3)) | ~ spl30_7), inference(backward_demodulation, [], [f210, f461])).
fof(f2636, plain, (~ spl30_3 | ~ spl30_19), inference(avatar_split_clause, [], [f2633, f510, f442])).
fof(f2633, plain, (~ (e2 = op(e3, e3)) | ~ spl30_19), inference(backward_demodulation, [], [f186, f512])).
fof(f512, plain, ((e2 = op(e2, e3)) | ~ spl30_19), inference(avatar_component_clause, [], [f510])).
fof(f186, plain, ~ (op(e2, e3) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f2632, plain, (~ spl30_21 | ~ spl30_22), inference(avatar_contradiction_clause, [], [f2631])).
fof(f2631, plain, ($false | (~ spl30_21 | ~ spl30_22)), inference(subsumption_resolution, [], [f2630, f211])).
fof(f2630, plain, ((e0 = e1) | (~ spl30_21 | ~ spl30_22)), inference(forward_demodulation, [], [f525, f521])).
fof(f2628, plain, (~ spl30_12 | ~ spl30_28), inference(avatar_split_clause, [], [f2626, f548, f480])).
fof(f2626, plain, (~ (e3 = op(e3, e1)) | ~ spl30_28), inference(backward_demodulation, [], [f174, f550])).
fof(f2627, plain, (~ spl30_20 | ~ spl30_28), inference(avatar_split_clause, [], [f2625, f548, f514])).
fof(f2625, plain, (~ (e3 = op(e2, e3)) | ~ spl30_28), inference(backward_demodulation, [], [f203, f550])).
fof(f2594, plain, (~ spl30_41 | ~ spl30_45), inference(avatar_split_clause, [], [f2586, f621, f604])).
fof(f2586, plain, (~ (e0 = op(e1, e1)) | ~ spl30_45), inference(backward_demodulation, [], [f193, f623])).
fof(f2584, plain, (~ spl30_62 | ~ spl30_63), inference(avatar_contradiction_clause, [], [f2583])).
fof(f2583, plain, ($false | (~ spl30_62 | ~ spl30_63)), inference(subsumption_resolution, [], [f2582, f214])).
fof(f2582, plain, ((e1 = e2) | (~ spl30_62 | ~ spl30_63)), inference(forward_demodulation, [], [f699, f695])).
fof(f2573, plain, (~ spl30_28 | ~ spl30_32), inference(avatar_split_clause, [], [f2369, f565, f548])).
fof(f2369, plain, (~ (e3 = op(e2, e1)) | ~ spl30_32), inference(backward_demodulation, [], [f199, f567])).
fof(f2536, plain, (~ spl30_1 | ~ spl30_13), inference(avatar_split_clause, [], [f2531, f485, f434])).
fof(f2531, plain, (~ (e0 = op(e3, e3)) | ~ spl30_13), inference(backward_demodulation, [], [f207, f487])).
fof(f2527, plain, (~ spl30_21 | ~ spl30_23), inference(avatar_contradiction_clause, [], [f2526])).
fof(f2526, plain, ($false | (~ spl30_21 | ~ spl30_23)), inference(subsumption_resolution, [], [f2525, f212])).
fof(f2525, plain, ((e0 = e2) | (~ spl30_21 | ~ spl30_23)), inference(backward_demodulation, [], [f529, f521])).
fof(f2497, plain, (~ spl30_41 | ~ spl30_56 | ~ spl30_59 | spl30_118), inference(avatar_contradiction_clause, [], [f2496])).
fof(f2496, plain, ($false | (~ spl30_41 | ~ spl30_56 | ~ spl30_59 | spl30_118)), inference(subsumption_resolution, [], [f2495, f669])).
fof(f2495, plain, (~ (e3 = op(e0, e2)) | (~ spl30_41 | ~ spl30_59 | spl30_118)), inference(forward_demodulation, [], [f2490, f682])).
fof(f2490, plain, (~ (e3 = op(e0, op(e0, e1))) | (~ spl30_41 | spl30_118)), inference(backward_demodulation, [], [f1150, f606])).
fof(f1150, plain, (~ (e3 = op(op(e1, e1), op(op(e1, e1), e1))) | spl30_118), inference(avatar_component_clause, [], [f1148])).
fof(f1148, plain, (spl30_118 <=> (e3 = op(op(e1, e1), op(op(e1, e1), e1)))), introduced(avatar_definition, [new_symbols(naming, [spl30_118])])).
fof(f2494, plain, (~ spl30_25 | ~ spl30_41), inference(avatar_split_clause, [], [f2485, f604, f536])).
fof(f2485, plain, (~ (e0 = op(e2, e1)) | ~ spl30_41), inference(backward_demodulation, [], [f172, f606])).
fof(f2482, plain, (~ spl30_43 | ~ spl30_47), inference(avatar_split_clause, [], [f2477, f629, f612])).
fof(f2477, plain, (~ (e2 = op(e1, e1)) | ~ spl30_47), inference(backward_demodulation, [], [f193, f631])).
fof(f2467, plain, (~ spl30_52 | ~ spl30_56), inference(avatar_split_clause, [], [f2464, f667, f650])).
fof(f2464, plain, (~ (e3 = op(e0, e3)) | ~ spl30_56), inference(backward_demodulation, [], [f192, f669])).
fof(f2461, plain, (~ spl30_55 | ~ spl30_59), inference(avatar_split_clause, [], [f2456, f680, f663])).
fof(f2456, plain, (~ (e2 = op(e0, e2)) | ~ spl30_59), inference(backward_demodulation, [], [f190, f682])).
fof(f2442, plain, (~ spl30_48 | ~ spl30_32), inference(avatar_split_clause, [], [f2441, f565, f633])).
fof(f2441, plain, (~ (e3 = op(e1, e0)) | ~ spl30_32), inference(forward_demodulation, [], [f166, f567])).
fof(f2437, plain, (~ spl30_47 | ~ spl30_62 | spl30_114), inference(avatar_split_clause, [], [f2187, f1128, f693, f629])).
fof(f2187, plain, (~ (e2 = op(e1, e0)) | (~ spl30_62 | spl30_114)), inference(forward_demodulation, [], [f1130, f695])).
fof(f2433, plain, (~ spl30_44 | ~ spl30_40), inference(avatar_split_clause, [], [f2432, f599, f616])).
fof(f2432, plain, (~ (e3 = op(e1, e1)) | ~ spl30_40), inference(forward_demodulation, [], [f196, f601])).
fof(f2431, plain, (~ spl30_24 | ~ spl30_32), inference(avatar_split_clause, [], [f2430, f565, f531])).
fof(f2430, plain, (~ (e3 = op(e2, e2)) | ~ spl30_32), inference(forward_demodulation, [], [f200, f567])).
fof(f2401, plain, (~ spl30_10 | ~ spl30_12), inference(avatar_contradiction_clause, [], [f2400])).
fof(f2400, plain, ($false | (~ spl30_10 | ~ spl30_12)), inference(subsumption_resolution, [], [f2399, f215])).
fof(f2399, plain, ((e1 = e3) | (~ spl30_10 | ~ spl30_12)), inference(backward_demodulation, [], [f482, f474])).
fof(f2397, plain, (spl30_5 | ~ spl30_15 | ~ spl30_69), inference(avatar_split_clause, [], [f2394, f757, f493, f451])).
fof(f2394, plain, ((e0 = op(e3, e2)) | (~ spl30_15 | ~ spl30_69)), inference(backward_demodulation, [], [f758, f495])).
fof(f2396, plain, (~ spl30_3 | ~ spl30_15), inference(avatar_split_clause, [], [f2392, f493, f442])).
fof(f2392, plain, (~ (e2 = op(e3, e3)) | ~ spl30_15), inference(backward_demodulation, [], [f207, f495])).
fof(f2384, plain, (~ spl30_22 | ~ spl30_24), inference(avatar_contradiction_clause, [], [f2383])).
fof(f2383, plain, ($false | (~ spl30_22 | ~ spl30_24)), inference(subsumption_resolution, [], [f2382, f215])).
fof(f2382, plain, ((e1 = e3) | (~ spl30_22 | ~ spl30_24)), inference(backward_demodulation, [], [f533, f525])).
fof(f2380, plain, (~ spl30_25 | ~ spl30_32 | ~ spl30_43 | spl30_118), inference(avatar_contradiction_clause, [], [f2379])).
fof(f2379, plain, ($false | (~ spl30_25 | ~ spl30_32 | ~ spl30_43 | spl30_118)), inference(subsumption_resolution, [], [f2375, f567])).
fof(f2375, plain, (~ (e3 = op(e2, e0)) | (~ spl30_25 | ~ spl30_43 | spl30_118)), inference(backward_demodulation, [], [f2243, f538])).
fof(f2243, plain, (~ (e3 = op(e2, op(e2, e1))) | (~ spl30_43 | spl30_118)), inference(backward_demodulation, [], [f1150, f614])).
fof(f2370, plain, (~ spl30_17 | ~ spl30_32 | spl30_76), inference(avatar_split_clause, [], [f2367, f803, f565, f502])).
fof(f803, plain, (spl30_76 <=> (e0 = op(e2, op(e2, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl30_76])])).
fof(f2367, plain, (~ (e0 = op(e2, e3)) | (~ spl30_32 | spl30_76)), inference(backward_demodulation, [], [f805, f567])).
fof(f805, plain, (~ (e0 = op(e2, op(e2, e0))) | spl30_76), inference(avatar_component_clause, [], [f803])).
fof(f2365, plain, (~ spl30_33 | ~ spl30_34), inference(avatar_contradiction_clause, [], [f2364])).
fof(f2364, plain, ($false | (~ spl30_33 | ~ spl30_34)), inference(subsumption_resolution, [], [f2363, f211])).
fof(f2363, plain, ((e0 = e1) | (~ spl30_33 | ~ spl30_34)), inference(backward_demodulation, [], [f576, f572])).
fof(f2359, plain, (spl30_50 | ~ spl30_60 | ~ spl30_91), inference(avatar_contradiction_clause, [], [f2358])).
fof(f2358, plain, ($false | (spl30_50 | ~ spl30_60 | ~ spl30_91)), inference(subsumption_resolution, [], [f2356, f643])).
fof(f643, plain, (~ (e1 = op(e0, e3)) | spl30_50), inference(avatar_component_clause, [], [f642])).
fof(f2356, plain, ((e1 = op(e0, e3)) | (~ spl30_60 | ~ spl30_91)), inference(backward_demodulation, [], [f911, f686])).
fof(f911, plain, ((e1 = op(e0, op(e0, e1))) | ~ spl30_91), inference(avatar_component_clause, [], [f910])).
fof(f2355, plain, (~ spl30_61 | ~ spl30_62), inference(avatar_contradiction_clause, [], [f2354])).
fof(f2354, plain, ($false | (~ spl30_61 | ~ spl30_62)), inference(subsumption_resolution, [], [f2353, f211])).
fof(f2353, plain, ((e0 = e1) | (~ spl30_61 | ~ spl30_62)), inference(backward_demodulation, [], [f695, f691])).
fof(f2340, plain, (~ spl30_60 | ~ spl30_52), inference(avatar_split_clause, [], [f2339, f650, f684])).
fof(f2339, plain, (~ (e3 = op(e0, e1)) | ~ spl30_52), inference(forward_demodulation, [], [f191, f652])).
fof(f2324, plain, (~ spl30_33 | ~ spl30_48 | spl30_84), inference(avatar_split_clause, [], [f2218, f859, f633, f570])).
fof(f2218, plain, (~ (e0 = op(e1, e3)) | (~ spl30_48 | spl30_84)), inference(forward_demodulation, [], [f861, f635])).
fof(f861, plain, (~ (e0 = op(e1, op(e1, e0))) | spl30_84), inference(avatar_component_clause, [], [f859])).
fof(f2314, plain, (~ spl30_20 | ~ spl30_52), inference(avatar_split_clause, [], [f2313, f650, f514])).
fof(f2313, plain, (~ (e3 = op(e2, e3)) | ~ spl30_52), inference(forward_demodulation, [], [f182, f652])).
fof(f2312, plain, (~ spl30_20 | ~ spl30_24), inference(avatar_split_clause, [], [f2311, f531, f514])).
fof(f2311, plain, (~ (e3 = op(e2, e3)) | ~ spl30_24), inference(forward_demodulation, [], [f204, f533])).
fof(f2304, plain, (~ spl30_5 | ~ spl30_24 | spl30_123), inference(avatar_split_clause, [], [f2272, f1174, f531, f451])).
fof(f2272, plain, (~ (e0 = op(e3, e2)) | (~ spl30_24 | spl30_123)), inference(backward_demodulation, [], [f1176, f533])).
fof(f2290, plain, (~ spl30_2 | ~ spl30_6), inference(avatar_split_clause, [], [f2287, f455, f438])).
fof(f2287, plain, (~ (e1 = op(e3, e3)) | ~ spl30_6), inference(backward_demodulation, [], [f210, f457])).
fof(f2283, plain, (~ spl30_8 | ~ spl30_12), inference(avatar_split_clause, [], [f2280, f480, f463])).
fof(f2280, plain, (~ (e3 = op(e3, e2)) | ~ spl30_12), inference(backward_demodulation, [], [f208, f482])).
fof(f2278, plain, (~ spl30_6 | ~ spl30_24 | spl30_126), inference(avatar_split_clause, [], [f2274, f1188, f531, f455])).
fof(f2274, plain, (~ (e1 = op(e3, e2)) | (~ spl30_24 | spl30_126)), inference(backward_demodulation, [], [f1190, f533])).
fof(f1190, plain, (~ (e1 = op(op(e2, e2), e2)) | spl30_126), inference(avatar_component_clause, [], [f1188])).
fof(f2267, plain, (~ spl30_26 | spl30_75), inference(avatar_contradiction_clause, [], [f2266])).
fof(f2266, plain, ($false | (~ spl30_26 | spl30_75)), inference(subsumption_resolution, [], [f2261, f542])).
fof(f2261, plain, (~ (e1 = op(e2, e1)) | (~ spl30_26 | spl30_75)), inference(backward_demodulation, [], [f800, f542])).
fof(f800, plain, (~ (e1 = op(e2, op(e2, e1))) | spl30_75), inference(avatar_component_clause, [], [f798])).
fof(f798, plain, (spl30_75 <=> (e1 = op(e2, op(e2, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl30_75])])).
fof(f2249, plain, (~ spl30_35 | ~ spl30_43), inference(avatar_split_clause, [], [f2240, f612, f578])).
fof(f2240, plain, (~ (e2 = op(e1, e3)) | ~ spl30_43), inference(backward_demodulation, [], [f197, f614])).
fof(f2248, plain, (~ spl30_11 | ~ spl30_43), inference(avatar_split_clause, [], [f2239, f612, f476])).
fof(f2239, plain, (~ (e2 = op(e3, e1)) | ~ spl30_43), inference(backward_demodulation, [], [f173, f614])).
fof(f2234, plain, (~ spl30_35 | ~ spl30_48 | ~ spl30_62 | spl30_113), inference(avatar_split_clause, [], [f2189, f1123, f693, f633, f578])).
fof(f1123, plain, (spl30_113 <=> (e2 = op(op(e0, e0), op(op(e0, e0), e0)))), introduced(avatar_definition, [new_symbols(naming, [spl30_113])])).
fof(f2189, plain, (~ (e2 = op(e1, e3)) | (~ spl30_48 | ~ spl30_62 | spl30_113)), inference(forward_demodulation, [], [f2188, f635])).
fof(f2188, plain, (~ (e2 = op(e1, op(e1, e0))) | (~ spl30_62 | spl30_113)), inference(forward_demodulation, [], [f1125, f695])).
fof(f1125, plain, (~ (e2 = op(op(e0, e0), op(op(e0, e0), e0))) | spl30_113), inference(avatar_component_clause, [], [f1123])).
fof(f2229, plain, (~ spl30_27 | ~ spl30_31), inference(avatar_split_clause, [], [f2228, f561, f544])).
fof(f2228, plain, (~ (e2 = op(e2, e1)) | ~ spl30_31), inference(forward_demodulation, [], [f199, f563])).
fof(f2214, plain, (~ spl30_22 | ~ spl30_37 | spl30_123), inference(avatar_contradiction_clause, [], [f2213])).
fof(f2213, plain, ($false | (~ spl30_22 | ~ spl30_37 | spl30_123)), inference(subsumption_resolution, [], [f2212, f589])).
fof(f2212, plain, (~ (e0 = op(e1, e2)) | (~ spl30_22 | spl30_123)), inference(forward_demodulation, [], [f1176, f525])).
fof(f2209, plain, (~ spl30_2 | ~ spl30_35 | ~ spl30_37 | spl30_133), inference(avatar_contradiction_clause, [], [f2208])).
fof(f2208, plain, ($false | (~ spl30_2 | ~ spl30_35 | ~ spl30_37 | spl30_133)), inference(subsumption_resolution, [], [f2207, f589])).
fof(f2207, plain, (~ (e0 = op(e1, e2)) | (~ spl30_2 | ~ spl30_35 | spl30_133)), inference(forward_demodulation, [], [f2206, f580])).
fof(f580, plain, ((e2 = op(e1, e3)) | ~ spl30_35), inference(avatar_component_clause, [], [f578])).
fof(f2206, plain, (~ (e0 = op(e1, op(e1, e3))) | (~ spl30_2 | spl30_133)), inference(forward_demodulation, [], [f1225, f440])).
fof(f1225, plain, (~ (e0 = op(op(e3, e3), op(op(e3, e3), e3))) | spl30_133), inference(avatar_component_clause, [], [f1223])).
fof(f1223, plain, (spl30_133 <=> (e0 = op(op(e3, e3), op(op(e3, e3), e3)))), introduced(avatar_definition, [new_symbols(naming, [spl30_133])])).
fof(f2186, plain, (~ spl30_2 | ~ spl30_35 | spl30_134), inference(avatar_contradiction_clause, [], [f2185])).
fof(f2185, plain, ($false | (~ spl30_2 | ~ spl30_35 | spl30_134)), inference(subsumption_resolution, [], [f2183, f580])).
fof(f2183, plain, (~ (e2 = op(e1, e3)) | (~ spl30_2 | spl30_134)), inference(backward_demodulation, [], [f1230, f440])).
fof(f1230, plain, (~ (e2 = op(op(e3, e3), e3)) | spl30_134), inference(avatar_component_clause, [], [f1228])).
fof(f2155, plain, (~ spl30_50 | ~ spl30_62), inference(avatar_split_clause, [], [f2147, f693, f642])).
fof(f2147, plain, (~ (e1 = op(e0, e3)) | ~ spl30_62), inference(backward_demodulation, [], [f189, f695])).
fof(f189, plain, ~ (op(e0, e0) = op(e0, e3)), inference(cnf_transformation, [], [f3])).
fof(f2140, plain, (spl30_62 | ~ spl30_57 | ~ spl30_91), inference(avatar_split_clause, [], [f2024, f910, f672, f693])).
fof(f2024, plain, ((op(e0, e0) = e1) | (~ spl30_57 | ~ spl30_91)), inference(forward_demodulation, [], [f911, f674])).
fof(f2136, plain, (spl30_47 | ~ spl30_37 | ~ spl30_82), inference(avatar_split_clause, [], [f2058, f849, f587, f629])).
fof(f849, plain, (spl30_82 <=> (e2 = op(e1, op(e1, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl30_82])])).
fof(f2058, plain, ((e2 = op(e1, e0)) | (~ spl30_37 | ~ spl30_82)), inference(backward_demodulation, [], [f850, f589])).
fof(f850, plain, ((e2 = op(e1, op(e1, e2))) | ~ spl30_82), inference(avatar_component_clause, [], [f849])).
fof(f2135, plain, (~ spl30_48 | ~ spl30_22 | ~ spl30_37 | spl30_124), inference(avatar_split_clause, [], [f2081, f1178, f587, f523, f633])).
fof(f1178, plain, (spl30_124 <=> (e3 = op(op(e2, e2), op(op(e2, e2), e2)))), introduced(avatar_definition, [new_symbols(naming, [spl30_124])])).
fof(f2081, plain, (~ (e3 = op(e1, e0)) | (~ spl30_22 | ~ spl30_37 | spl30_124)), inference(forward_demodulation, [], [f2076, f589])).
fof(f2076, plain, (~ (e3 = op(e1, op(e1, e2))) | (~ spl30_22 | spl30_124)), inference(backward_demodulation, [], [f1180, f525])).
fof(f1180, plain, (~ (e3 = op(op(e2, e2), op(op(e2, e2), e2))) | spl30_124), inference(avatar_component_clause, [], [f1178])).
fof(f2117, plain, (~ spl30_51 | ~ spl30_35), inference(avatar_split_clause, [], [f2116, f578, f646])).
fof(f2116, plain, (~ (e2 = op(e0, e3)) | ~ spl30_35), inference(forward_demodulation, [], [f181, f580])).
fof(f2113, plain, (~ spl30_27 | ~ spl30_22 | spl30_74), inference(avatar_split_clause, [], [f2075, f793, f523, f544])).
fof(f2075, plain, (~ (e2 = op(e2, e1)) | (~ spl30_22 | spl30_74)), inference(backward_demodulation, [], [f795, f525])).
fof(f2103, plain, (~ spl30_1 | ~ spl30_17), inference(avatar_split_clause, [], [f2102, f502, f434])).
fof(f2102, plain, (~ (e0 = op(e3, e3)) | ~ spl30_17), inference(forward_demodulation, [], [f186, f504])).
fof(f2099, plain, (~ spl30_13 | ~ spl30_14), inference(avatar_contradiction_clause, [], [f2098])).
fof(f2098, plain, ($false | (~ spl30_13 | ~ spl30_14)), inference(subsumption_resolution, [], [f2097, f211])).
fof(f2097, plain, ((e0 = e1) | (~ spl30_13 | ~ spl30_14)), inference(backward_demodulation, [], [f491, f487])).
fof(f2079, plain, (~ spl30_18 | ~ spl30_22), inference(avatar_split_clause, [], [f2072, f523, f506])).
fof(f2072, plain, (~ (e1 = op(e2, e3)) | ~ spl30_22), inference(backward_demodulation, [], [f204, f525])).
fof(f2071, plain, (~ spl30_26 | ~ spl30_28), inference(avatar_contradiction_clause, [], [f2070])).
fof(f2070, plain, ($false | (~ spl30_26 | ~ spl30_28)), inference(subsumption_resolution, [], [f2069, f215])).
fof(f2069, plain, ((e1 = e3) | (~ spl30_26 | ~ spl30_28)), inference(backward_demodulation, [], [f550, f542])).
fof(f2067, plain, (~ spl30_35 | ~ spl30_36), inference(avatar_contradiction_clause, [], [f2066])).
fof(f2066, plain, ($false | (~ spl30_35 | ~ spl30_36)), inference(subsumption_resolution, [], [f2065, f216])).
fof(f216, plain, ~ (e2 = e3), inference(cnf_transformation, [], [f4])).
fof(f2065, plain, ((e2 = e3) | (~ spl30_35 | ~ spl30_36)), inference(backward_demodulation, [], [f584, f580])).
fof(f2064, plain, (~ spl30_21 | ~ spl30_37), inference(avatar_split_clause, [], [f2060, f587, f519])).
fof(f2060, plain, (~ (e0 = op(e2, e2)) | ~ spl30_37), inference(backward_demodulation, [], [f178, f589])).
fof(f2055, plain, (~ spl30_18 | ~ spl30_50), inference(avatar_split_clause, [], [f2052, f642, f506])).
fof(f2052, plain, (~ (e1 = op(e2, e3)) | ~ spl30_50), inference(backward_demodulation, [], [f182, f644])).
fof(f2018, plain, (~ spl30_51 | ~ spl30_1 | spl30_134), inference(avatar_split_clause, [], [f1990, f1228, f434, f646])).
fof(f1990, plain, (~ (e2 = op(e0, e3)) | (~ spl30_1 | spl30_134)), inference(backward_demodulation, [], [f1230, f436])).
fof(f2014, plain, (~ spl30_38 | ~ spl30_46), inference(avatar_split_clause, [], [f2013, f625, f591])).
fof(f2013, plain, (~ (e1 = op(e1, e2)) | ~ spl30_46), inference(forward_demodulation, [], [f194, f627])).
fof(f1998, plain, (~ spl30_15 | ~ spl30_31), inference(avatar_split_clause, [], [f1997, f561, f493])).
fof(f1997, plain, (~ (e2 = op(e3, e0)) | ~ spl30_31), inference(forward_demodulation, [], [f168, f563])).
fof(f1982, plain, (~ spl30_10 | spl30_68), inference(avatar_contradiction_clause, [], [f1981])).
fof(f1981, plain, ($false | (~ spl30_10 | spl30_68)), inference(subsumption_resolution, [], [f1978, f474])).
fof(f1978, plain, (~ (e1 = op(e3, e1)) | (~ spl30_10 | spl30_68)), inference(backward_demodulation, [], [f754, f474])).
fof(f1976, plain, (~ spl30_2 | ~ spl30_18), inference(avatar_split_clause, [], [f1973, f506, f438])).
fof(f1973, plain, (~ (e1 = op(e3, e3)) | ~ spl30_18), inference(backward_demodulation, [], [f186, f508])).
fof(f1975, plain, (~ spl30_18 | ~ spl30_28 | spl30_73), inference(avatar_contradiction_clause, [], [f1974])).
fof(f1974, plain, ($false | (~ spl30_18 | ~ spl30_28 | spl30_73)), inference(subsumption_resolution, [], [f1971, f550])).
fof(f1971, plain, (~ (e3 = op(e2, e1)) | (~ spl30_18 | spl30_73)), inference(backward_demodulation, [], [f790, f508])).
fof(f790, plain, (~ (e3 = op(e2, op(e2, e3))) | spl30_73), inference(avatar_component_clause, [], [f788])).
fof(f788, plain, (spl30_73 <=> (e3 = op(e2, op(e2, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl30_73])])).
fof(f1966, plain, (~ spl30_21 | ~ spl30_31 | spl30_74), inference(avatar_contradiction_clause, [], [f1965])).
fof(f1965, plain, ($false | (~ spl30_21 | ~ spl30_31 | spl30_74)), inference(subsumption_resolution, [], [f1960, f563])).
fof(f1960, plain, (~ (e2 = op(e2, e0)) | (~ spl30_21 | spl30_74)), inference(backward_demodulation, [], [f795, f521])).
fof(f1956, plain, (~ spl30_22 | ~ spl30_38), inference(avatar_split_clause, [], [f1955, f591, f523])).
fof(f1955, plain, (~ (e1 = op(e2, e2)) | ~ spl30_38), inference(backward_demodulation, [], [f178, f593])).
fof(f1952, plain, (~ spl30_45 | ~ spl30_46), inference(avatar_contradiction_clause, [], [f1951])).
fof(f1951, plain, ($false | (~ spl30_45 | ~ spl30_46)), inference(subsumption_resolution, [], [f1950, f211])).
fof(f1950, plain, ((e0 = e1) | (~ spl30_45 | ~ spl30_46)), inference(forward_demodulation, [], [f627, f623])).
fof(f1949, plain, (~ spl30_57 | ~ spl30_58), inference(avatar_contradiction_clause, [], [f1948])).
fof(f1948, plain, ($false | (~ spl30_57 | ~ spl30_58)), inference(subsumption_resolution, [], [f1947, f211])).
fof(f1947, plain, ((e0 = e1) | (~ spl30_57 | ~ spl30_58)), inference(backward_demodulation, [], [f678, f674])).
fof(f1946, plain, (~ spl30_58 | ~ spl30_59), inference(avatar_contradiction_clause, [], [f1945])).
fof(f1945, plain, ($false | (~ spl30_58 | ~ spl30_59)), inference(subsumption_resolution, [], [f1944, f214])).
fof(f1944, plain, ((e1 = e2) | (~ spl30_58 | ~ spl30_59)), inference(backward_demodulation, [], [f682, f678])).
fof(f1904, plain, (~ spl30_17 | ~ spl30_28 | ~ spl30_43 | spl30_121), inference(avatar_split_clause, [], [f1903, f1163, f612, f548, f502])).
fof(f1163, plain, (spl30_121 <=> (e0 = op(op(e1, e1), op(op(e1, e1), e1)))), introduced(avatar_definition, [new_symbols(naming, [spl30_121])])).
fof(f1903, plain, (~ (e0 = op(e2, e3)) | (~ spl30_28 | ~ spl30_43 | spl30_121)), inference(forward_demodulation, [], [f1902, f550])).
fof(f1902, plain, (~ (e0 = op(e2, op(e2, e1))) | (~ spl30_43 | spl30_121)), inference(forward_demodulation, [], [f1165, f614])).
fof(f1165, plain, (~ (e0 = op(op(e1, e1), op(op(e1, e1), e1))) | spl30_121), inference(avatar_component_clause, [], [f1163])).
fof(f1898, plain, (~ spl30_21 | ~ spl30_31 | spl30_76), inference(avatar_split_clause, [], [f1897, f803, f561, f519])).
fof(f1897, plain, (~ (e0 = op(e2, e2)) | (~ spl30_31 | spl30_76)), inference(forward_demodulation, [], [f805, f563])).
fof(f1856, plain, (~ spl30_27 | ~ spl30_28), inference(avatar_contradiction_clause, [], [f1855])).
fof(f1855, plain, ($false | (~ spl30_27 | ~ spl30_28)), inference(subsumption_resolution, [], [f1854, f216])).
fof(f1854, plain, ((e2 = e3) | (~ spl30_27 | ~ spl30_28)), inference(forward_demodulation, [], [f550, f546])).
fof(f1853, plain, (~ spl30_30 | ~ spl30_31), inference(avatar_contradiction_clause, [], [f1852])).
fof(f1852, plain, ($false | (~ spl30_30 | ~ spl30_31)), inference(subsumption_resolution, [], [f1851, f214])).
fof(f1851, plain, ((e1 = e2) | (~ spl30_30 | ~ spl30_31)), inference(forward_demodulation, [], [f563, f559])).
fof(f1842, plain, (~ spl30_42 | ~ spl30_43), inference(avatar_contradiction_clause, [], [f1841])).
fof(f1841, plain, ($false | (~ spl30_42 | ~ spl30_43)), inference(subsumption_resolution, [], [f1840, f214])).
fof(f1840, plain, ((e1 = e2) | (~ spl30_42 | ~ spl30_43)), inference(forward_demodulation, [], [f614, f610])).
fof(f1839, plain, (~ spl30_49 | ~ spl30_51), inference(avatar_contradiction_clause, [], [f1838])).
fof(f1838, plain, ($false | (~ spl30_49 | ~ spl30_51)), inference(subsumption_resolution, [], [f1837, f212])).
fof(f1837, plain, ((e0 = e2) | (~ spl30_49 | ~ spl30_51)), inference(forward_demodulation, [], [f648, f640])).
fof(f1836, plain, (~ spl30_54 | ~ spl30_55), inference(avatar_contradiction_clause, [], [f1835])).
fof(f1835, plain, ($false | (~ spl30_54 | ~ spl30_55)), inference(subsumption_resolution, [], [f1834, f214])).
fof(f1834, plain, ((e1 = e2) | (~ spl30_54 | ~ spl30_55)), inference(forward_demodulation, [], [f665, f661])).
fof(f1821, plain, (spl30_64 | ~ spl30_49 | ~ spl30_89), inference(avatar_split_clause, [], [f1820, f900, f638, f701])).
fof(f900, plain, (spl30_89 <=> (e3 = op(e0, op(e0, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl30_89])])).
fof(f1820, plain, ((op(e0, e0) = e3) | (~ spl30_49 | ~ spl30_89)), inference(forward_demodulation, [], [f901, f640])).
fof(f901, plain, ((e3 = op(e0, op(e0, e3))) | ~ spl30_89), inference(avatar_component_clause, [], [f900])).
fof(f1811, plain, (~ spl30_27 | ~ spl30_59), inference(avatar_contradiction_clause, [], [f1810])).
fof(f1810, plain, ($false | (~ spl30_27 | ~ spl30_59)), inference(subsumption_resolution, [], [f1809, f682])).
fof(f1809, plain, (~ (e2 = op(e0, e1)) | ~ spl30_27), inference(forward_demodulation, [], [f170, f546])).
fof(f1802, plain, (~ spl30_22 | ~ spl30_30), inference(avatar_split_clause, [], [f1542, f557, f523])).
fof(f1542, plain, (~ (e1 = op(e2, e2)) | ~ spl30_30), inference(backward_demodulation, [], [f200, f559])).
fof(f1801, plain, (~ spl30_23 | ~ spl30_27), inference(avatar_split_clause, [], [f1800, f544, f527])).
fof(f1800, plain, (~ (e2 = op(e2, e2)) | ~ spl30_27), inference(forward_demodulation, [], [f202, f546])).
fof(f1797, plain, (~ spl30_22 | ~ spl30_27 | spl30_75), inference(avatar_split_clause, [], [f1769, f798, f544, f523])).
fof(f1769, plain, (~ (e1 = op(e2, e2)) | (~ spl30_27 | spl30_75)), inference(backward_demodulation, [], [f800, f546])).
fof(f1790, plain, (~ spl30_13 | ~ spl30_9), inference(avatar_split_clause, [], [f1789, f468, f485])).
fof(f1789, plain, (~ (e0 = op(e3, e0)) | ~ spl30_9), inference(forward_demodulation, [], [f205, f470])).
fof(f1786, plain, (~ spl30_63 | ~ spl30_15), inference(avatar_split_clause, [], [f1785, f493, f697])).
fof(f1785, plain, (~ (op(e0, e0) = e2) | ~ spl30_15), inference(forward_demodulation, [], [f165, f495])).
fof(f1781, plain, (~ spl30_14 | ~ spl30_9 | spl30_68), inference(avatar_split_clause, [], [f1777, f752, f468, f489])).
fof(f1777, plain, (~ (e1 = op(e3, e0)) | (~ spl30_9 | spl30_68)), inference(backward_demodulation, [], [f754, f470])).
fof(f1775, plain, (~ spl30_13 | ~ spl30_15), inference(avatar_contradiction_clause, [], [f1774])).
fof(f1774, plain, ($false | (~ spl30_13 | ~ spl30_15)), inference(subsumption_resolution, [], [f1773, f212])).
fof(f1773, plain, ((e0 = e2) | (~ spl30_13 | ~ spl30_15)), inference(forward_demodulation, [], [f495, f487])).
fof(f1772, plain, (~ spl30_11 | ~ spl30_27), inference(avatar_split_clause, [], [f1770, f544, f476])).
fof(f1770, plain, (~ (e2 = op(e3, e1)) | ~ spl30_27), inference(backward_demodulation, [], [f174, f546])).
fof(f1767, plain, (~ spl30_30 | ~ spl30_32), inference(avatar_contradiction_clause, [], [f1766])).
fof(f1766, plain, ($false | (~ spl30_30 | ~ spl30_32)), inference(subsumption_resolution, [], [f1765, f215])).
fof(f1765, plain, ((e1 = e3) | (~ spl30_30 | ~ spl30_32)), inference(forward_demodulation, [], [f567, f559])).
fof(f1764, plain, (~ spl30_39 | spl30_82), inference(avatar_contradiction_clause, [], [f1763])).
fof(f1763, plain, ($false | (~ spl30_39 | spl30_82)), inference(subsumption_resolution, [], [f1762, f597])).
fof(f1762, plain, (~ (e2 = op(e1, e2)) | (~ spl30_39 | spl30_82)), inference(backward_demodulation, [], [f851, f597])).
fof(f851, plain, (~ (e2 = op(e1, op(e1, e2))) | spl30_82), inference(avatar_component_clause, [], [f849])).
fof(f1757, plain, (~ spl30_38 | ~ spl30_42), inference(avatar_split_clause, [], [f1751, f608, f591])).
fof(f1751, plain, (~ (e1 = op(e1, e2)) | ~ spl30_42), inference(backward_demodulation, [], [f196, f610])).
fof(f1745, plain, (~ spl30_45 | spl30_84), inference(avatar_contradiction_clause, [], [f1744])).
fof(f1744, plain, ($false | (~ spl30_45 | spl30_84)), inference(subsumption_resolution, [], [f1739, f623])).
fof(f1739, plain, (~ (e0 = op(e1, e0)) | (~ spl30_45 | spl30_84)), inference(backward_demodulation, [], [f861, f623])).
fof(f1725, plain, (~ spl30_59 | ~ spl30_60), inference(avatar_contradiction_clause, [], [f1724])).
fof(f1724, plain, ($false | (~ spl30_59 | ~ spl30_60)), inference(subsumption_resolution, [], [f1723, f216])).
fof(f1723, plain, ((e2 = e3) | (~ spl30_59 | ~ spl30_60)), inference(backward_demodulation, [], [f686, f682])).
fof(f1697, plain, (~ spl30_54 | ~ spl30_21 | spl30_126), inference(avatar_split_clause, [], [f1657, f1188, f519, f659])).
fof(f1657, plain, (~ (e1 = op(e0, e2)) | (~ spl30_21 | spl30_126)), inference(backward_demodulation, [], [f1190, f521])).
fof(f1679, plain, (~ spl30_40 | ~ spl30_8), inference(avatar_split_clause, [], [f1678, f463, f599])).
fof(f1678, plain, (~ (e3 = op(e1, e2)) | ~ spl30_8), inference(forward_demodulation, [], [f179, f465])).
fof(f465, plain, ((e3 = op(e3, e2)) | ~ spl30_8), inference(avatar_component_clause, [], [f463])).
fof(f1668, plain, (~ spl30_13 | spl30_69), inference(avatar_contradiction_clause, [], [f1667])).
fof(f1667, plain, ($false | (~ spl30_13 | spl30_69)), inference(subsumption_resolution, [], [f1663, f487])).
fof(f1663, plain, (~ (e0 = op(e3, e0)) | (~ spl30_13 | spl30_69)), inference(backward_demodulation, [], [f759, f487])).
fof(f1639, plain, (~ spl30_49 | ~ spl30_50), inference(avatar_contradiction_clause, [], [f1638])).
fof(f1638, plain, ($false | (~ spl30_49 | ~ spl30_50)), inference(subsumption_resolution, [], [f1637, f211])).
fof(f1637, plain, ((e0 = e1) | (~ spl30_49 | ~ spl30_50)), inference(backward_demodulation, [], [f644, f640])).
fof(f1635, plain, (~ spl30_55 | spl30_90), inference(avatar_contradiction_clause, [], [f1634])).
fof(f1634, plain, ($false | (~ spl30_55 | spl30_90)), inference(subsumption_resolution, [], [f1632, f665])).
fof(f1632, plain, (~ (e2 = op(e0, e2)) | (~ spl30_55 | spl30_90)), inference(backward_demodulation, [], [f907, f665])).
fof(f1633, plain, (~ spl30_23 | ~ spl30_55), inference(avatar_split_clause, [], [f1631, f663, f527])).
fof(f1631, plain, (~ (e2 = op(e2, e2)) | ~ spl30_55), inference(backward_demodulation, [], [f176, f665])).
fof(f1630, plain, (~ spl30_58 | spl30_91), inference(avatar_contradiction_clause, [], [f1629])).
fof(f1629, plain, ($false | (~ spl30_58 | spl30_91)), inference(subsumption_resolution, [], [f1626, f678])).
fof(f1626, plain, (~ (e1 = op(e0, e1)) | (~ spl30_58 | spl30_91)), inference(backward_demodulation, [], [f912, f678])).
fof(f1628, plain, (~ spl30_54 | ~ spl30_58), inference(avatar_split_clause, [], [f1625, f676, f659])).
fof(f1625, plain, (~ (e1 = op(e0, e2)) | ~ spl30_58), inference(backward_demodulation, [], [f190, f678])).
fof(f1613, plain, (~ spl30_60 | ~ spl30_50 | spl30_89), inference(avatar_split_clause, [], [f1612, f900, f642, f684])).
fof(f1612, plain, (~ (e3 = op(e0, e1)) | (~ spl30_50 | spl30_89)), inference(forward_demodulation, [], [f902, f644])).
fof(f902, plain, (~ (e3 = op(e0, op(e0, e3))) | spl30_89), inference(avatar_component_clause, [], [f900])).
fof(f1598, plain, (~ spl30_2 | ~ spl30_50), inference(avatar_contradiction_clause, [], [f1597])).
fof(f1597, plain, ($false | (~ spl30_2 | ~ spl30_50)), inference(subsumption_resolution, [], [f1596, f644])).
fof(f1596, plain, (~ (e1 = op(e0, e3)) | ~ spl30_2), inference(forward_demodulation, [], [f183, f440])).
fof(f1588, plain, (~ spl30_35 | ~ spl30_19), inference(avatar_split_clause, [], [f1584, f510, f578])).
fof(f1584, plain, (~ (e2 = op(e1, e3)) | ~ spl30_19), inference(backward_demodulation, [], [f184, f512])).
fof(f1586, plain, (~ spl30_24 | ~ spl30_19 | spl30_73), inference(avatar_split_clause, [], [f1582, f788, f510, f531])).
fof(f1582, plain, (~ (e3 = op(e2, e2)) | (~ spl30_19 | spl30_73)), inference(backward_demodulation, [], [f790, f512])).
fof(f1579, plain, (~ spl30_14 | ~ spl30_2), inference(avatar_split_clause, [], [f1438, f438, f489])).
fof(f1438, plain, (~ (e1 = op(e3, e0)) | ~ spl30_2), inference(forward_demodulation, [], [f207, f440])).
fof(f1576, plain, (~ spl30_56 | ~ spl30_8), inference(avatar_split_clause, [], [f1573, f463, f667])).
fof(f1573, plain, (~ (e3 = op(e0, e2)) | ~ spl30_8), inference(backward_demodulation, [], [f177, f465])).
fof(f1574, plain, (~ spl30_16 | ~ spl30_8), inference(avatar_split_clause, [], [f1571, f463, f497])).
fof(f1571, plain, (~ (e3 = op(e3, e0)) | ~ spl30_8), inference(backward_demodulation, [], [f206, f465])).
fof(f1568, plain, (~ spl30_6 | ~ spl30_38), inference(avatar_split_clause, [], [f1567, f591, f455])).
fof(f1567, plain, (~ (e1 = op(e3, e2)) | ~ spl30_38), inference(forward_demodulation, [], [f179, f593])).
fof(f1565, plain, (~ spl30_5 | ~ spl30_6), inference(avatar_contradiction_clause, [], [f1564])).
fof(f1564, plain, ($false | (~ spl30_5 | ~ spl30_6)), inference(subsumption_resolution, [], [f1563, f211])).
fof(f1563, plain, ((e0 = e1) | (~ spl30_5 | ~ spl30_6)), inference(backward_demodulation, [], [f457, f453])).
fof(f1555, plain, (~ spl30_23 | spl30_74), inference(avatar_contradiction_clause, [], [f1554])).
fof(f1554, plain, ($false | (~ spl30_23 | spl30_74)), inference(subsumption_resolution, [], [f1551, f529])).
fof(f1551, plain, (~ (e2 = op(e2, e2)) | (~ spl30_23 | spl30_74)), inference(backward_demodulation, [], [f795, f529])).
fof(f1546, plain, (~ spl30_25 | ~ spl30_30 | spl30_76), inference(avatar_split_clause, [], [f1543, f803, f557, f536])).
fof(f1543, plain, (~ (e0 = op(e2, e1)) | (~ spl30_30 | spl30_76)), inference(backward_demodulation, [], [f805, f559])).
fof(f1528, plain, (~ spl30_42 | spl30_83), inference(avatar_contradiction_clause, [], [f1527])).
fof(f1527, plain, ($false | (~ spl30_42 | spl30_83)), inference(subsumption_resolution, [], [f1522, f610])).
fof(f1522, plain, (~ (e1 = op(e1, e1)) | (~ spl30_42 | spl30_83)), inference(backward_demodulation, [], [f856, f610])).
fof(f1517, plain, (~ spl30_45 | ~ spl30_47), inference(avatar_contradiction_clause, [], [f1516])).
fof(f1516, plain, ($false | (~ spl30_45 | ~ spl30_47)), inference(subsumption_resolution, [], [f1515, f212])).
fof(f1515, plain, ((e0 = e2) | (~ spl30_45 | ~ spl30_47)), inference(backward_demodulation, [], [f631, f623])).
fof(f1505, plain, (~ spl30_30 | ~ spl30_63 | spl30_111), inference(avatar_split_clause, [], [f1500, f1114, f697, f557])).
fof(f1500, plain, (~ (e1 = op(e2, e0)) | (~ spl30_63 | spl30_111)), inference(backward_demodulation, [], [f1116, f699])).
fof(f1116, plain, (~ (e1 = op(op(e0, e0), e0)) | spl30_111), inference(avatar_component_clause, [], [f1114])).
fof(f1504, plain, (~ spl30_53 | ~ spl30_63 | spl30_92), inference(avatar_split_clause, [], [f1499, f915, f697, f655])).
fof(f1499, plain, (~ (e0 = op(e0, e2)) | (~ spl30_63 | spl30_92)), inference(backward_demodulation, [], [f917, f699])).
fof(f1503, plain, (~ spl30_55 | ~ spl30_63), inference(avatar_split_clause, [], [f1497, f697, f663])).
fof(f1497, plain, (~ (e2 = op(e0, e2)) | ~ spl30_63), inference(backward_demodulation, [], [f188, f699])).
fof(f1491, plain, (~ spl30_63 | ~ spl30_47), inference(avatar_split_clause, [], [f1490, f629, f697])).
fof(f1490, plain, (~ (op(e0, e0) = e2) | ~ spl30_47), inference(forward_demodulation, [], [f163, f631])).
fof(f1485, plain, (~ spl30_62 | ~ spl30_57 | spl30_91), inference(avatar_split_clause, [], [f1256, f910, f672, f693])).
fof(f1256, plain, (~ (op(e0, e0) = e1) | (~ spl30_57 | spl30_91)), inference(backward_demodulation, [], [f912, f674])).
fof(f1482, plain, (~ spl30_54 | ~ spl30_50), inference(avatar_split_clause, [], [f1481, f642, f659])).
fof(f1481, plain, (~ (e1 = op(e0, e2)) | ~ spl30_50), inference(forward_demodulation, [], [f192, f644])).
fof(f1480, plain, (~ spl30_50 | ~ spl30_52), inference(avatar_contradiction_clause, [], [f1479])).
fof(f1479, plain, ($false | (~ spl30_50 | ~ spl30_52)), inference(subsumption_resolution, [], [f1478, f215])).
fof(f1478, plain, ((e1 = e3) | (~ spl30_50 | ~ spl30_52)), inference(forward_demodulation, [], [f652, f644])).
fof(f1448, plain, (~ spl30_17 | ~ spl30_19), inference(avatar_contradiction_clause, [], [f1447])).
fof(f1447, plain, ($false | (~ spl30_17 | ~ spl30_19)), inference(subsumption_resolution, [], [f1446, f212])).
fof(f1446, plain, ((e0 = e2) | (~ spl30_17 | ~ spl30_19)), inference(forward_demodulation, [], [f512, f504])).
fof(f1443, plain, (~ spl30_6 | ~ spl30_14), inference(avatar_contradiction_clause, [], [f1442])).
fof(f1442, plain, ($false | (~ spl30_6 | ~ spl30_14)), inference(subsumption_resolution, [], [f1441, f491])).
fof(f1441, plain, (~ (e1 = op(e3, e0)) | ~ spl30_6), inference(forward_demodulation, [], [f206, f457])).
fof(f1435, plain, (~ spl30_10 | ~ spl30_6), inference(avatar_split_clause, [], [f1434, f455, f472])).
fof(f1434, plain, (~ (e1 = op(e3, e1)) | ~ spl30_6), inference(forward_demodulation, [], [f208, f457])).
fof(f1427, plain, (~ spl30_6 | ~ spl30_8), inference(avatar_contradiction_clause, [], [f1426])).
fof(f1426, plain, ($false | (~ spl30_6 | ~ spl30_8)), inference(subsumption_resolution, [], [f1425, f215])).
fof(f1425, plain, ((e1 = e3) | (~ spl30_6 | ~ spl30_8)), inference(forward_demodulation, [], [f465, f457])).
fof(f1423, plain, (~ spl30_33 | ~ spl30_2 | spl30_129), inference(avatar_split_clause, [], [f1420, f1204, f438, f570])).
fof(f1420, plain, (~ (e0 = op(e1, e3)) | (~ spl30_2 | spl30_129)), inference(backward_demodulation, [], [f1206, f440])).
fof(f1206, plain, (~ (e0 = op(op(e3, e3), e3)) | spl30_129), inference(avatar_component_clause, [], [f1204])).
fof(f1417, plain, (~ spl30_6 | ~ spl30_7), inference(avatar_contradiction_clause, [], [f1416])).
fof(f1416, plain, ($false | (~ spl30_6 | ~ spl30_7)), inference(subsumption_resolution, [], [f1415, f214])).
fof(f1415, plain, ((e1 = e2) | (~ spl30_6 | ~ spl30_7)), inference(backward_demodulation, [], [f461, f457])).
fof(f1414, plain, (~ spl30_7 | ~ spl30_8), inference(avatar_contradiction_clause, [], [f1413])).
fof(f1413, plain, ($false | (~ spl30_7 | ~ spl30_8)), inference(subsumption_resolution, [], [f1412, f216])).
fof(f1412, plain, ((e2 = e3) | (~ spl30_7 | ~ spl30_8)), inference(backward_demodulation, [], [f465, f461])).
fof(f1411, plain, (~ spl30_3 | ~ spl30_8 | spl30_67), inference(avatar_split_clause, [], [f1409, f747, f463, f442])).
fof(f1409, plain, (~ (e2 = op(e3, e3)) | (~ spl30_8 | spl30_67)), inference(backward_demodulation, [], [f749, f465])).
fof(f1410, plain, (~ spl30_4 | ~ spl30_8), inference(avatar_split_clause, [], [f1408, f463, f446])).
fof(f1408, plain, (~ (e3 = op(e3, e3)) | ~ spl30_8), inference(backward_demodulation, [], [f210, f465])).
fof(f1407, plain, (~ spl30_9 | ~ spl30_10), inference(avatar_contradiction_clause, [], [f1406])).
fof(f1406, plain, ($false | (~ spl30_9 | ~ spl30_10)), inference(subsumption_resolution, [], [f1405, f211])).
fof(f1405, plain, ((e0 = e1) | (~ spl30_9 | ~ spl30_10)), inference(backward_demodulation, [], [f474, f470])).
fof(f1404, plain, (~ spl30_10 | ~ spl30_11), inference(avatar_contradiction_clause, [], [f1403])).
fof(f1403, plain, ($false | (~ spl30_10 | ~ spl30_11)), inference(subsumption_resolution, [], [f1402, f214])).
fof(f1402, plain, ((e1 = e2) | (~ spl30_10 | ~ spl30_11)), inference(backward_demodulation, [], [f478, f474])).
fof(f1401, plain, (~ spl30_11 | ~ spl30_12), inference(avatar_contradiction_clause, [], [f1400])).
fof(f1400, plain, ($false | (~ spl30_11 | ~ spl30_12)), inference(subsumption_resolution, [], [f1399, f216])).
fof(f1399, plain, ((e2 = e3) | (~ spl30_11 | ~ spl30_12)), inference(backward_demodulation, [], [f482, f478])).
fof(f1397, plain, (~ spl30_4 | ~ spl30_12), inference(avatar_split_clause, [], [f1394, f480, f446])).
fof(f1394, plain, (~ (e3 = op(e3, e3)) | ~ spl30_12), inference(backward_demodulation, [], [f209, f482])).
fof(f1392, plain, (~ spl30_14 | ~ spl30_16), inference(avatar_contradiction_clause, [], [f1391])).
fof(f1391, plain, ($false | (~ spl30_14 | ~ spl30_16)), inference(subsumption_resolution, [], [f1390, f215])).
fof(f1390, plain, ((e1 = e3) | (~ spl30_14 | ~ spl30_16)), inference(backward_demodulation, [], [f499, f491])).
fof(f1389, plain, (~ spl30_1 | ~ spl30_16 | spl30_69), inference(avatar_split_clause, [], [f1385, f757, f497, f434])).
fof(f1385, plain, (~ (e0 = op(e3, e3)) | (~ spl30_16 | spl30_69)), inference(backward_demodulation, [], [f759, f499])).
fof(f1381, plain, (~ spl30_17 | ~ spl30_18), inference(avatar_contradiction_clause, [], [f1380])).
fof(f1380, plain, ($false | (~ spl30_17 | ~ spl30_18)), inference(subsumption_resolution, [], [f1379, f211])).
fof(f1379, plain, ((e0 = e1) | (~ spl30_17 | ~ spl30_18)), inference(backward_demodulation, [], [f508, f504])).
fof(f1378, plain, (~ spl30_18 | ~ spl30_19), inference(avatar_contradiction_clause, [], [f1377])).
fof(f1377, plain, ($false | (~ spl30_18 | ~ spl30_19)), inference(subsumption_resolution, [], [f1376, f214])).
fof(f1376, plain, ((e1 = e2) | (~ spl30_18 | ~ spl30_19)), inference(backward_demodulation, [], [f512, f508])).
fof(f1375, plain, (~ spl30_19 | ~ spl30_20), inference(avatar_contradiction_clause, [], [f1374])).
fof(f1374, plain, ($false | (~ spl30_19 | ~ spl30_20)), inference(subsumption_resolution, [], [f1373, f216])).
fof(f1373, plain, ((e2 = e3) | (~ spl30_19 | ~ spl30_20)), inference(backward_demodulation, [], [f516, f512])).
fof(f1372, plain, (~ spl30_20 | spl30_73), inference(avatar_contradiction_clause, [], [f1371])).
fof(f1371, plain, ($false | (~ spl30_20 | spl30_73)), inference(subsumption_resolution, [], [f1369, f516])).
fof(f1369, plain, (~ (e3 = op(e2, e3)) | (~ spl30_20 | spl30_73)), inference(backward_demodulation, [], [f790, f516])).
fof(f1370, plain, (~ spl30_4 | ~ spl30_20), inference(avatar_split_clause, [], [f1368, f514, f446])).
fof(f1368, plain, (~ (e3 = op(e3, e3)) | ~ spl30_20), inference(backward_demodulation, [], [f186, f516])).
fof(f1353, plain, (~ spl30_21 | ~ spl30_25), inference(avatar_split_clause, [], [f1349, f536, f519])).
fof(f1349, plain, (~ (e0 = op(e2, e2)) | ~ spl30_25), inference(backward_demodulation, [], [f202, f538])).
fof(f1347, plain, (~ spl30_29 | spl30_76), inference(avatar_contradiction_clause, [], [f1346])).
fof(f1346, plain, ($false | (~ spl30_29 | spl30_76)), inference(subsumption_resolution, [], [f1341, f555])).
fof(f1341, plain, (~ (e0 = op(e2, e0)) | (~ spl30_29 | spl30_76)), inference(backward_demodulation, [], [f805, f555])).
fof(f1343, plain, (~ spl30_25 | ~ spl30_29), inference(avatar_split_clause, [], [f1338, f553, f536])).
fof(f1338, plain, (~ (e0 = op(e2, e1)) | ~ spl30_29), inference(backward_demodulation, [], [f199, f555])).
fof(f1342, plain, (~ spl30_13 | ~ spl30_29), inference(avatar_split_clause, [], [f1337, f553, f485])).
fof(f1337, plain, (~ (e0 = op(e3, e0)) | ~ spl30_29), inference(backward_demodulation, [], [f168, f555])).
fof(f1336, plain, (~ spl30_34 | ~ spl30_36), inference(avatar_contradiction_clause, [], [f1335])).
fof(f1335, plain, ($false | (~ spl30_34 | ~ spl30_36)), inference(subsumption_resolution, [], [f1334, f215])).
fof(f1334, plain, ((e1 = e3) | (~ spl30_34 | ~ spl30_36)), inference(backward_demodulation, [], [f584, f576])).
fof(f1333, plain, (~ spl30_36 | spl30_81), inference(avatar_contradiction_clause, [], [f1332])).
fof(f1332, plain, ($false | (~ spl30_36 | spl30_81)), inference(subsumption_resolution, [], [f1329, f584])).
fof(f1329, plain, (~ (e3 = op(e1, e3)) | (~ spl30_36 | spl30_81)), inference(backward_demodulation, [], [f846, f584])).
fof(f846, plain, (~ (e3 = op(e1, op(e1, e3))) | spl30_81), inference(avatar_component_clause, [], [f844])).
fof(f844, plain, (spl30_81 <=> (e3 = op(e1, op(e1, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl30_81])])).
fof(f1331, plain, (~ spl30_4 | ~ spl30_36), inference(avatar_split_clause, [], [f1328, f582, f446])).
fof(f1328, plain, (~ (e3 = op(e3, e3)) | ~ spl30_36), inference(backward_demodulation, [], [f185, f584])).
fof(f1326, plain, (~ spl30_37 | ~ spl30_47 | spl30_82), inference(avatar_contradiction_clause, [], [f1325])).
fof(f1325, plain, ($false | (~ spl30_37 | ~ spl30_47 | spl30_82)), inference(subsumption_resolution, [], [f1321, f631])).
fof(f1321, plain, (~ (e2 = op(e1, e0)) | (~ spl30_37 | spl30_82)), inference(backward_demodulation, [], [f851, f589])).
fof(f1324, plain, (~ spl30_33 | ~ spl30_37), inference(avatar_split_clause, [], [f1320, f587, f570])).
fof(f1320, plain, (~ (e0 = op(e1, e3)) | ~ spl30_37), inference(backward_demodulation, [], [f198, f589])).
fof(f1313, plain, (~ spl30_46 | ~ spl30_41 | spl30_83), inference(avatar_split_clause, [], [f1306, f854, f604, f625])).
fof(f1306, plain, (~ (e1 = op(e1, e0)) | (~ spl30_41 | spl30_83)), inference(backward_demodulation, [], [f856, f606])).
fof(f1312, plain, (~ spl30_33 | ~ spl30_41), inference(avatar_split_clause, [], [f1303, f604, f570])).
fof(f1303, plain, (~ (e0 = op(e1, e3)) | ~ spl30_41), inference(backward_demodulation, [], [f197, f606])).
fof(f1310, plain, (~ spl30_9 | ~ spl30_41), inference(avatar_split_clause, [], [f1301, f604, f468])).
fof(f1301, plain, (~ (e0 = op(e3, e1)) | ~ spl30_41), inference(backward_demodulation, [], [f173, f606])).
fof(f1299, plain, (~ spl30_37 | ~ spl30_47 | spl30_84), inference(avatar_split_clause, [], [f1293, f859, f629, f587])).
fof(f1293, plain, (~ (e0 = op(e1, e2)) | (~ spl30_47 | spl30_84)), inference(backward_demodulation, [], [f861, f631])).
fof(f1294, plain, (~ spl30_31 | ~ spl30_47), inference(avatar_split_clause, [], [f1288, f629, f561])).
fof(f1288, plain, (~ (e2 = op(e2, e0)) | ~ spl30_47), inference(backward_demodulation, [], [f166, f631])).
fof(f1287, plain, (~ spl30_50 | ~ spl30_51), inference(avatar_contradiction_clause, [], [f1286])).
fof(f1286, plain, ($false | (~ spl30_50 | ~ spl30_51)), inference(subsumption_resolution, [], [f1285, f214])).
fof(f1285, plain, ((e1 = e2) | (~ spl30_50 | ~ spl30_51)), inference(backward_demodulation, [], [f648, f644])).
fof(f1284, plain, (~ spl30_51 | ~ spl30_52), inference(avatar_contradiction_clause, [], [f1283])).
fof(f1283, plain, ($false | (~ spl30_51 | ~ spl30_52)), inference(subsumption_resolution, [], [f1282, f216])).
fof(f1282, plain, ((e2 = e3) | (~ spl30_51 | ~ spl30_52)), inference(backward_demodulation, [], [f652, f648])).
fof(f1281, plain, (~ spl30_52 | spl30_89), inference(avatar_contradiction_clause, [], [f1280])).
fof(f1280, plain, ($false | (~ spl30_52 | spl30_89)), inference(subsumption_resolution, [], [f1276, f652])).
fof(f1276, plain, (~ (e3 = op(e0, e3)) | (~ spl30_52 | spl30_89)), inference(backward_demodulation, [], [f902, f652])).
fof(f1277, plain, (~ spl30_36 | ~ spl30_52), inference(avatar_split_clause, [], [f1273, f650, f582])).
fof(f1273, plain, (~ (e3 = op(e1, e3)) | ~ spl30_52), inference(backward_demodulation, [], [f181, f652])).
fof(f1271, plain, (~ spl30_49 | ~ spl30_53), inference(avatar_split_clause, [], [f1266, f655, f638])).
fof(f1266, plain, (~ (e0 = op(e0, e3)) | ~ spl30_53), inference(backward_demodulation, [], [f192, f657])).
fof(f1270, plain, (~ spl30_5 | ~ spl30_53), inference(avatar_split_clause, [], [f1265, f655, f451])).
fof(f1265, plain, (~ (e0 = op(e3, e2)) | ~ spl30_53), inference(backward_demodulation, [], [f177, f657])).
fof(f1268, plain, (~ spl30_37 | ~ spl30_53), inference(avatar_split_clause, [], [f1263, f655, f587])).
fof(f1263, plain, (~ (e0 = op(e1, e2)) | ~ spl30_53), inference(backward_demodulation, [], [f175, f657])).
fof(f1257, plain, (~ spl30_41 | ~ spl30_57), inference(avatar_split_clause, [], [f1251, f672, f604])).
fof(f1251, plain, (~ (e0 = op(e1, e1)) | ~ spl30_57), inference(backward_demodulation, [], [f169, f674])).
fof(f1250, plain, (~ spl30_61 | spl30_92), inference(avatar_contradiction_clause, [], [f1249])).
fof(f1249, plain, ($false | (~ spl30_61 | spl30_92)), inference(subsumption_resolution, [], [f1242, f691])).
fof(f1242, plain, (~ (e0 = op(e0, e0)) | (~ spl30_61 | spl30_92)), inference(backward_demodulation, [], [f917, f691])).
fof(f1248, plain, (~ spl30_49 | ~ spl30_61), inference(avatar_split_clause, [], [f1238, f689, f638])).
fof(f1238, plain, (~ (e0 = op(e0, e3)) | ~ spl30_61), inference(backward_demodulation, [], [f189, f691])).
fof(f1247, plain, (~ spl30_53 | ~ spl30_61), inference(avatar_split_clause, [], [f1237, f689, f655])).
fof(f1237, plain, (~ (e0 = op(e0, e2)) | ~ spl30_61), inference(backward_demodulation, [], [f188, f691])).
fof(f1245, plain, (~ spl30_13 | ~ spl30_61), inference(avatar_split_clause, [], [f1235, f689, f485])).
fof(f1235, plain, (~ (e0 = op(e3, e0)) | ~ spl30_61), inference(backward_demodulation, [], [f165, f691])).
fof(f1244, plain, (~ spl30_29 | ~ spl30_61), inference(avatar_split_clause, [], [f1234, f689, f553])).
fof(f1234, plain, (~ (e0 = op(e2, e0)) | ~ spl30_61), inference(backward_demodulation, [], [f164, f691])).
fof(f1232, plain, (~ spl30_134 | ~ spl30_2 | ~ spl30_133), inference(avatar_split_clause, [], [f432, f1223, f438, f1228])).
fof(f432, plain, (~ (e0 = op(op(e3, e3), op(op(e3, e3), e3))) | ~ (e1 = op(e3, e3)) | ~ (e2 = op(op(e3, e3), e3))), inference(cnf_transformation, [], [f53])).
fof(f53, plain, (~ (e0 = op(op(e3, e3), op(op(e3, e3), e3))) | ~ (e1 = op(e3, e3)) | ~ (e2 = op(op(e3, e3), e3))), inference(ennf_transformation, [], [f29])).
fof(f29, plain, ~ ((e0 = op(op(e3, e3), op(op(e3, e3), e3))) & (e1 = op(e3, e3)) & (e2 = op(op(e3, e3), e3))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG159+1.p', ax29)).
fof(f1231, plain, (~ spl30_134 | ~ spl30_1 | ~ spl30_131), inference(avatar_split_clause, [], [f431, f1213, f434, f1228])).
fof(f431, plain, (~ (e1 = op(op(e3, e3), op(op(e3, e3), e3))) | ~ (e0 = op(e3, e3)) | ~ (e2 = op(op(e3, e3), e3))), inference(cnf_transformation, [], [f52])).
fof(f52, plain, (~ (e1 = op(op(e3, e3), op(op(e3, e3), e3))) | ~ (e0 = op(e3, e3)) | ~ (e2 = op(op(e3, e3), e3))), inference(ennf_transformation, [], [f28])).
fof(f28, plain, ~ ((e1 = op(op(e3, e3), op(op(e3, e3), e3))) & (e0 = op(e3, e3)) & (e2 = op(op(e3, e3), e3))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG159+1.p', ax28)).
fof(f1211, plain, (~ spl30_129 | ~ spl30_2 | ~ spl30_130), inference(avatar_split_clause, [], [f427, f1208, f438, f1204])).
fof(f427, plain, (~ (e2 = op(op(e3, e3), op(op(e3, e3), e3))) | ~ (e1 = op(e3, e3)) | ~ (e0 = op(op(e3, e3), e3))), inference(cnf_transformation, [], [f48])).
fof(f48, plain, (~ (e2 = op(op(e3, e3), op(op(e3, e3), e3))) | ~ (e1 = op(e3, e3)) | ~ (e0 = op(op(e3, e3), e3))), inference(ennf_transformation, [], [f24])).
fof(f24, plain, ~ ((e2 = op(op(e3, e3), op(op(e3, e3), e3))) & (e1 = op(e3, e3)) & (e0 = op(op(e3, e3), e3))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG159+1.p', ax24)).
fof(f1196, plain, (~ spl30_126 | ~ spl30_24 | ~ spl30_127), inference(avatar_split_clause, [], [f424, f1193, f531, f1188])).
fof(f424, plain, (~ (e0 = op(op(e2, e2), op(op(e2, e2), e2))) | ~ (e3 = op(e2, e2)) | ~ (e1 = op(op(e2, e2), e2))), inference(cnf_transformation, [], [f45])).
fof(f45, plain, (~ (e0 = op(op(e2, e2), op(op(e2, e2), e2))) | ~ (e3 = op(e2, e2)) | ~ (e1 = op(op(e2, e2), e2))), inference(ennf_transformation, [], [f21])).
fof(f21, plain, ~ ((e0 = op(op(e2, e2), op(op(e2, e2), e2))) & (e3 = op(e2, e2)) & (e1 = op(op(e2, e2), e2))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG159+1.p', ax21)).
fof(f1186, plain, (~ spl30_123 | ~ spl30_24 | ~ spl30_125), inference(avatar_split_clause, [], [f422, f1183, f531, f1174])).
fof(f422, plain, (~ (e1 = op(op(e2, e2), op(op(e2, e2), e2))) | ~ (e3 = op(e2, e2)) | ~ (e0 = op(op(e2, e2), e2))), inference(cnf_transformation, [], [f43])).
fof(f43, plain, (~ (e1 = op(op(e2, e2), op(op(e2, e2), e2))) | ~ (e3 = op(e2, e2)) | ~ (e0 = op(op(e2, e2), e2))), inference(ennf_transformation, [], [f19])).
fof(f19, plain, ~ ((e1 = op(op(e2, e2), op(op(e2, e2), e2))) & (e3 = op(e2, e2)) & (e0 = op(op(e2, e2), e2))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG159+1.p', ax19)).
fof(f1181, plain, (~ spl30_123 | ~ spl30_22 | ~ spl30_124), inference(avatar_split_clause, [], [f421, f1178, f523, f1174])).
fof(f421, plain, (~ (e3 = op(op(e2, e2), op(op(e2, e2), e2))) | ~ (e1 = op(e2, e2)) | ~ (e0 = op(op(e2, e2), e2))), inference(cnf_transformation, [], [f42])).
fof(f42, plain, (~ (e3 = op(op(e2, e2), op(op(e2, e2), e2))) | ~ (e1 = op(e2, e2)) | ~ (e0 = op(op(e2, e2), e2))), inference(ennf_transformation, [], [f18])).
fof(f18, plain, ~ ((e3 = op(op(e2, e2), op(op(e2, e2), e2))) & (e1 = op(e2, e2)) & (e0 = op(op(e2, e2), e2))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG159+1.p', ax18)).
fof(f1172, plain, (~ spl30_122 | ~ spl30_43 | ~ spl30_121), inference(avatar_split_clause, [], [f420, f1163, f612, f1168])).
fof(f420, plain, (~ (e0 = op(op(e1, e1), op(op(e1, e1), e1))) | ~ (e2 = op(e1, e1)) | ~ (e3 = op(op(e1, e1), e1))), inference(cnf_transformation, [], [f41])).
fof(f41, plain, (~ (e0 = op(op(e1, e1), op(op(e1, e1), e1))) | ~ (e2 = op(e1, e1)) | ~ (e3 = op(op(e1, e1), e1))), inference(ennf_transformation, [], [f17])).
fof(f17, plain, ~ ((e0 = op(op(e1, e1), op(op(e1, e1), e1))) & (e2 = op(e1, e1)) & (e3 = op(op(e1, e1), e1))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG159+1.p', ax17)).
fof(f1161, plain, (~ spl30_120 | ~ spl30_41 | ~ spl30_118), inference(avatar_split_clause, [], [f417, f1148, f604, f1158])).
fof(f417, plain, (~ (e3 = op(op(e1, e1), op(op(e1, e1), e1))) | ~ (e0 = op(e1, e1)) | ~ (e2 = op(op(e1, e1), e1))), inference(cnf_transformation, [], [f38])).
fof(f38, plain, (~ (e3 = op(op(e1, e1), op(op(e1, e1), e1))) | ~ (e0 = op(e1, e1)) | ~ (e2 = op(op(e1, e1), e1))), inference(ennf_transformation, [], [f14])).
fof(f14, plain, ~ ((e3 = op(op(e1, e1), op(op(e1, e1), e1))) & (e0 = op(e1, e1)) & (e2 = op(op(e1, e1), e1))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG159+1.p', ax14)).
fof(f1156, plain, (~ spl30_117 | ~ spl30_44 | ~ spl30_119), inference(avatar_split_clause, [], [f416, f1153, f616, f1144])).
fof(f416, plain, (~ (e2 = op(op(e1, e1), op(op(e1, e1), e1))) | ~ (e3 = op(e1, e1)) | ~ (e0 = op(op(e1, e1), e1))), inference(cnf_transformation, [], [f37])).
fof(f37, plain, (~ (e2 = op(op(e1, e1), op(op(e1, e1), e1))) | ~ (e3 = op(e1, e1)) | ~ (e0 = op(op(e1, e1), e1))), inference(ennf_transformation, [], [f13])).
fof(f13, plain, ~ ((e2 = op(op(e1, e1), op(op(e1, e1), e1))) & (e3 = op(e1, e1)) & (e0 = op(op(e1, e1), e1))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG159+1.p', ax13)).
fof(f1151, plain, (~ spl30_117 | ~ spl30_43 | ~ spl30_118), inference(avatar_split_clause, [], [f415, f1148, f612, f1144])).
fof(f415, plain, (~ (e3 = op(op(e1, e1), op(op(e1, e1), e1))) | ~ (e2 = op(e1, e1)) | ~ (e0 = op(op(e1, e1), e1))), inference(cnf_transformation, [], [f36])).
fof(f36, plain, (~ (e3 = op(op(e1, e1), op(op(e1, e1), e1))) | ~ (e2 = op(e1, e1)) | ~ (e0 = op(op(e1, e1), e1))), inference(ennf_transformation, [], [f12])).
fof(f12, plain, ~ ((e3 = op(op(e1, e1), op(op(e1, e1), e1))) & (e2 = op(e1, e1)) & (e0 = op(op(e1, e1), e1))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG159+1.p', ax12)).
fof(f1142, plain, (~ spl30_116 | ~ spl30_63 | ~ spl30_115), inference(avatar_split_clause, [], [f414, f1133, f697, f1138])).
fof(f414, plain, (~ (e1 = op(op(e0, e0), op(op(e0, e0), e0))) | ~ (op(e0, e0) = e2) | ~ (e3 = op(op(e0, e0), e0))), inference(cnf_transformation, [], [f35])).
fof(f35, plain, (~ (e1 = op(op(e0, e0), op(op(e0, e0), e0))) | ~ (op(e0, e0) = e2) | ~ (e3 = op(op(e0, e0), e0))), inference(ennf_transformation, [], [f11])).
fof(f11, plain, ~ ((e1 = op(op(e0, e0), op(op(e0, e0), e0))) & (op(e0, e0) = e2) & (e3 = op(op(e0, e0), e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG159+1.p', ax11)).
fof(f1141, plain, (~ spl30_116 | ~ spl30_62 | ~ spl30_113), inference(avatar_split_clause, [], [f413, f1123, f693, f1138])).
fof(f413, plain, (~ (e2 = op(op(e0, e0), op(op(e0, e0), e0))) | ~ (op(e0, e0) = e1) | ~ (e3 = op(op(e0, e0), e0))), inference(cnf_transformation, [], [f34])).
fof(f34, plain, (~ (e2 = op(op(e0, e0), op(op(e0, e0), e0))) | ~ (op(e0, e0) = e1) | ~ (e3 = op(op(e0, e0), e0))), inference(ennf_transformation, [], [f10])).
fof(f10, plain, ~ ((e2 = op(op(e0, e0), op(op(e0, e0), e0))) & (op(e0, e0) = e1) & (e3 = op(op(e0, e0), e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG159+1.p', ax10)).
fof(f1136, plain, (~ spl30_114 | ~ spl30_64 | ~ spl30_115), inference(avatar_split_clause, [], [f412, f1133, f701, f1128])).
fof(f412, plain, (~ (e1 = op(op(e0, e0), op(op(e0, e0), e0))) | ~ (op(e0, e0) = e3) | ~ (e2 = op(op(e0, e0), e0))), inference(cnf_transformation, [], [f33])).
fof(f33, plain, (~ (e1 = op(op(e0, e0), op(op(e0, e0), e0))) | ~ (op(e0, e0) = e3) | ~ (e2 = op(op(e0, e0), e0))), inference(ennf_transformation, [], [f9])).
fof(f9, plain, ~ ((e1 = op(op(e0, e0), op(op(e0, e0), e0))) & (op(e0, e0) = e3) & (e2 = op(op(e0, e0), e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG159+1.p', ax9)).
fof(f1131, plain, (~ spl30_114 | ~ spl30_62 | ~ spl30_112), inference(avatar_split_clause, [], [f411, f1118, f693, f1128])).
fof(f411, plain, (~ (e3 = op(op(e0, e0), op(op(e0, e0), e0))) | ~ (op(e0, e0) = e1) | ~ (e2 = op(op(e0, e0), e0))), inference(cnf_transformation, [], [f32])).
fof(f32, plain, (~ (e3 = op(op(e0, e0), op(op(e0, e0), e0))) | ~ (op(e0, e0) = e1) | ~ (e2 = op(op(e0, e0), e0))), inference(ennf_transformation, [], [f8])).
fof(f8, plain, ~ ((e3 = op(op(e0, e0), op(op(e0, e0), e0))) & (op(e0, e0) = e1) & (e2 = op(op(e0, e0), e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG159+1.p', ax8)).
fof(f1112, plain, (spl30_110 | spl30_109 | spl30_108 | spl30_107 | spl30_106 | spl30_105 | spl30_104 | spl30_103 | spl30_102 | spl30_101 | spl30_100 | spl30_99 | spl30_98 | spl30_97 | spl30_96 | spl30_4), inference(avatar_split_clause, [], [f397, f446, f952, f962, f972, f982, f992, f1002, f1012, f1022, f1032, f1042, f1052, f1062, f1072, f1082, f1092])).
fof(f1092, plain, (spl30_110 <=> sP0), introduced(avatar_definition, [new_symbols(naming, [spl30_110])])).
fof(f1082, plain, (spl30_109 <=> sP1), introduced(avatar_definition, [new_symbols(naming, [spl30_109])])).
fof(f1072, plain, (spl30_108 <=> sP2), introduced(avatar_definition, [new_symbols(naming, [spl30_108])])).
fof(f1062, plain, (spl30_107 <=> sP3), introduced(avatar_definition, [new_symbols(naming, [spl30_107])])).
fof(f1052, plain, (spl30_106 <=> sP4), introduced(avatar_definition, [new_symbols(naming, [spl30_106])])).
fof(f1042, plain, (spl30_105 <=> sP5), introduced(avatar_definition, [new_symbols(naming, [spl30_105])])).
fof(f1032, plain, (spl30_104 <=> sP6), introduced(avatar_definition, [new_symbols(naming, [spl30_104])])).
fof(f1022, plain, (spl30_103 <=> sP7), introduced(avatar_definition, [new_symbols(naming, [spl30_103])])).
fof(f1012, plain, (spl30_102 <=> sP8), introduced(avatar_definition, [new_symbols(naming, [spl30_102])])).
fof(f1002, plain, (spl30_101 <=> sP9), introduced(avatar_definition, [new_symbols(naming, [spl30_101])])).
fof(f992, plain, (spl30_100 <=> sP10), introduced(avatar_definition, [new_symbols(naming, [spl30_100])])).
fof(f982, plain, (spl30_99 <=> sP11), introduced(avatar_definition, [new_symbols(naming, [spl30_99])])).
fof(f972, plain, (spl30_98 <=> sP12), introduced(avatar_definition, [new_symbols(naming, [spl30_98])])).
fof(f962, plain, (spl30_97 <=> sP13), introduced(avatar_definition, [new_symbols(naming, [spl30_97])])).
fof(f952, plain, (spl30_96 <=> sP14), introduced(avatar_definition, [new_symbols(naming, [spl30_96])])).
fof(f397, plain, ((e3 = op(e3, e3)) | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f84])).
fof(f84, plain, (((~ (e3 = op(e3, op(e3, e3))) & ~ (e2 = op(e3, op(e3, e2))) & ~ (e1 = op(e3, op(e3, e1))) & ~ (e0 = op(e3, op(e3, e0))) & (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15) & ((~ (e3 = op(e3, op(e3, e3))) & ~ (e2 = op(e3, op(e3, e2))) & ~ (e1 = op(e3, op(e3, e1))) & ~ (e0 = op(e3, op(e3, e0))) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0)), inference(definition_folding, [], [f5, e83, e82, e81, e80, e79, e78, e77, e76, e75, e74, e73, e72, e71, e70, e69, e68, e67, e66, e65, e64, e63, e62, e61, e60, e59, e58, e57, e56, e55, e54])).
fof(f54, plain, ((~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP0), inference(usedef, [], [e54])).
fof(e54, plain, (sP0 <=> (~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f55, plain, ((~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | ~ sP1), inference(usedef, [], [e55])).
fof(e55, plain, (sP1 <=> (~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f56, plain, ((~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | ~ sP2), inference(usedef, [], [e56])).
fof(e56, plain, (sP2 <=> (~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f57, plain, ((~ (e3 = op(e3, op(e3, e3))) & ~ (e2 = op(e3, op(e3, e2))) & ~ (e1 = op(e3, op(e3, e1))) & ~ (e0 = op(e3, op(e3, e0))) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | ~ sP3), inference(usedef, [], [e57])).
fof(e57, plain, (sP3 <=> (~ (e3 = op(e3, op(e3, e3))) & ~ (e2 = op(e3, op(e3, e2))) & ~ (e1 = op(e3, op(e3, e1))) & ~ (e0 = op(e3, op(e3, e0))) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f58, plain, ((~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | ~ sP4), inference(usedef, [], [e58])).
fof(e58, plain, (sP4 <=> (~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f59, plain, ((~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP5), inference(usedef, [], [e59])).
fof(e59, plain, (sP5 <=> (~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f60, plain, ((~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | ~ sP6), inference(usedef, [], [e60])).
fof(e60, plain, (sP6 <=> (~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f61, plain, ((~ (e3 = op(e3, op(e3, e3))) & ~ (e2 = op(e3, op(e3, e2))) & ~ (e1 = op(e3, op(e3, e1))) & ~ (e0 = op(e3, op(e3, e0))) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | ~ sP7), inference(usedef, [], [e61])).
fof(e61, plain, (sP7 <=> (~ (e3 = op(e3, op(e3, e3))) & ~ (e2 = op(e3, op(e3, e2))) & ~ (e1 = op(e3, op(e3, e1))) & ~ (e0 = op(e3, op(e3, e0))) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f62, plain, ((~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | ~ sP8), inference(usedef, [], [e62])).
fof(e62, plain, (sP8 <=> (~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f63, plain, ((~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | ~ sP9), inference(usedef, [], [e63])).
fof(e63, plain, (sP9 <=> (~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f64, plain, ((~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP10), inference(usedef, [], [e64])).
fof(e64, plain, (sP10 <=> (~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f65, plain, ((~ (e3 = op(e3, op(e3, e3))) & ~ (e2 = op(e3, op(e3, e2))) & ~ (e1 = op(e3, op(e3, e1))) & ~ (e0 = op(e3, op(e3, e0))) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | ~ sP11), inference(usedef, [], [e65])).
fof(e65, plain, (sP11 <=> (~ (e3 = op(e3, op(e3, e3))) & ~ (e2 = op(e3, op(e3, e2))) & ~ (e1 = op(e3, op(e3, e1))) & ~ (e0 = op(e3, op(e3, e0))) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f66, plain, ((~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | ~ sP12), inference(usedef, [], [e66])).
fof(e66, plain, (sP12 <=> (~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f67, plain, ((~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | ~ sP13), inference(usedef, [], [e67])).
fof(e67, plain, (sP13 <=> (~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f68, plain, ((~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | ~ sP14), inference(usedef, [], [e68])).
fof(e68, plain, (sP14 <=> (~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f69, plain, ((~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP15), inference(usedef, [], [e69])).
fof(e69, plain, (sP15 <=> (~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & (e0 = op(e0, e0)) & (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f70, plain, ((~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & (e0 = op(e1, e0)) & (e0 = op(e0, e1))) | ~ sP16), inference(usedef, [], [e70])).
fof(e70, plain, (sP16 <=> (~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & (e0 = op(e1, e0)) & (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f71, plain, ((~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & (e0 = op(e2, e0)) & (e0 = op(e0, e2))) | ~ sP17), inference(usedef, [], [e71])).
fof(e71, plain, (sP17 <=> (~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & (e0 = op(e2, e0)) & (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f72, plain, ((~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & (e0 = op(e3, e0)) & (e0 = op(e0, e3))) | ~ sP18), inference(usedef, [], [e72])).
fof(e72, plain, (sP18 <=> (~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & (e0 = op(e3, e0)) & (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f73, plain, ((~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & (e1 = op(e0, e1)) & (e1 = op(e1, e0))) | ~ sP19), inference(usedef, [], [e73])).
fof(e73, plain, (sP19 <=> (~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & (e1 = op(e0, e1)) & (e1 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP19])])).
fof(f74, plain, ((~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP20), inference(usedef, [], [e74])).
fof(e74, plain, (sP20 <=> (~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & (e1 = op(e1, e1)) & (e1 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP20])])).
fof(f75, plain, ((~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & (e1 = op(e2, e1)) & (e1 = op(e1, e2))) | ~ sP21), inference(usedef, [], [e75])).
fof(e75, plain, (sP21 <=> (~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & (e1 = op(e2, e1)) & (e1 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP21])])).
fof(f76, plain, ((~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & (e1 = op(e3, e1)) & (e1 = op(e1, e3))) | ~ sP22), inference(usedef, [], [e76])).
fof(e76, plain, (sP22 <=> (~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & (e1 = op(e3, e1)) & (e1 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP22])])).
fof(f77, plain, ((~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & (e2 = op(e0, e2)) & (e2 = op(e2, e0))) | ~ sP23), inference(usedef, [], [e77])).
fof(e77, plain, (sP23 <=> (~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & (e2 = op(e0, e2)) & (e2 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP23])])).
fof(f78, plain, ((~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & (e2 = op(e1, e2)) & (e2 = op(e2, e1))) | ~ sP24), inference(usedef, [], [e78])).
fof(e78, plain, (sP24 <=> (~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & (e2 = op(e1, e2)) & (e2 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP24])])).
fof(f79, plain, ((~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP25), inference(usedef, [], [e79])).
fof(e79, plain, (sP25 <=> (~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & (e2 = op(e2, e2)) & (e2 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP25])])).
fof(f80, plain, ((~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & (e2 = op(e3, e2)) & (e2 = op(e2, e3))) | ~ sP26), inference(usedef, [], [e80])).
fof(e80, plain, (sP26 <=> (~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & (e2 = op(e3, e2)) & (e2 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP26])])).
fof(f81, plain, ((~ (e3 = op(e3, op(e3, e3))) & ~ (e2 = op(e3, op(e3, e2))) & ~ (e1 = op(e3, op(e3, e1))) & ~ (e0 = op(e3, op(e3, e0))) & (e3 = op(e0, e3)) & (e3 = op(e3, e0))) | ~ sP27), inference(usedef, [], [e81])).
fof(e81, plain, (sP27 <=> (~ (e3 = op(e3, op(e3, e3))) & ~ (e2 = op(e3, op(e3, e2))) & ~ (e1 = op(e3, op(e3, e1))) & ~ (e0 = op(e3, op(e3, e0))) & (e3 = op(e0, e3)) & (e3 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP27])])).
fof(f82, plain, ((~ (e3 = op(e3, op(e3, e3))) & ~ (e2 = op(e3, op(e3, e2))) & ~ (e1 = op(e3, op(e3, e1))) & ~ (e0 = op(e3, op(e3, e0))) & (e3 = op(e1, e3)) & (e3 = op(e3, e1))) | ~ sP28), inference(usedef, [], [e82])).
fof(e82, plain, (sP28 <=> (~ (e3 = op(e3, op(e3, e3))) & ~ (e2 = op(e3, op(e3, e2))) & ~ (e1 = op(e3, op(e3, e1))) & ~ (e0 = op(e3, op(e3, e0))) & (e3 = op(e1, e3)) & (e3 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP28])])).
fof(f83, plain, ((~ (e3 = op(e3, op(e3, e3))) & ~ (e2 = op(e3, op(e3, e2))) & ~ (e1 = op(e3, op(e3, e1))) & ~ (e0 = op(e3, op(e3, e0))) & (e3 = op(e2, e3)) & (e3 = op(e3, e2))) | ~ sP29), inference(usedef, [], [e83])).
fof(e83, plain, (sP29 <=> (~ (e3 = op(e3, op(e3, e3))) & ~ (e2 = op(e3, op(e3, e2))) & ~ (e1 = op(e3, op(e3, e1))) & ~ (e0 = op(e3, op(e3, e0))) & (e3 = op(e2, e3)) & (e3 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP29])])).
fof(f5, plain, (((~ (e3 = op(e3, op(e3, e3))) & ~ (e2 = op(e3, op(e3, e2))) & ~ (e1 = op(e3, op(e3, e1))) & ~ (e0 = op(e3, op(e3, e0))) & (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | (~ (e3 = op(e3, op(e3, e3))) & ~ (e2 = op(e3, op(e3, e2))) & ~ (e1 = op(e3, op(e3, e1))) & ~ (e0 = op(e3, op(e3, e0))) & (e3 = op(e2, e3)) & (e3 = op(e3, e2))) | (~ (e3 = op(e3, op(e3, e3))) & ~ (e2 = op(e3, op(e3, e2))) & ~ (e1 = op(e3, op(e3, e1))) & ~ (e0 = op(e3, op(e3, e0))) & (e3 = op(e1, e3)) & (e3 = op(e3, e1))) | (~ (e3 = op(e3, op(e3, e3))) & ~ (e2 = op(e3, op(e3, e2))) & ~ (e1 = op(e3, op(e3, e1))) & ~ (e0 = op(e3, op(e3, e0))) & (e3 = op(e0, e3)) & (e3 = op(e3, e0))) | (~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & (e2 = op(e3, e2)) & (e2 = op(e2, e3))) | (~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | (~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & (e2 = op(e1, e2)) & (e2 = op(e2, e1))) | (~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & (e2 = op(e0, e2)) & (e2 = op(e2, e0))) | (~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & (e1 = op(e3, e1)) & (e1 = op(e1, e3))) | (~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & (e1 = op(e2, e1)) & (e1 = op(e1, e2))) | (~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | (~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & (e1 = op(e0, e1)) & (e1 = op(e1, e0))) | (~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & (e0 = op(e3, e0)) & (e0 = op(e0, e3))) | (~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & (e0 = op(e2, e0)) & (e0 = op(e0, e2))) | (~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & (e0 = op(e1, e0)) & (e0 = op(e0, e1))) | (~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & (e0 = op(e0, e0)) & (e0 = op(e0, e0)))) & ((~ (e3 = op(e3, op(e3, e3))) & ~ (e2 = op(e3, op(e3, e2))) & ~ (e1 = op(e3, op(e3, e1))) & ~ (e0 = op(e3, op(e3, e0))) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | (~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | (~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | (~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | (~ (e3 = op(e3, op(e3, e3))) & ~ (e2 = op(e3, op(e3, e2))) & ~ (e1 = op(e3, op(e3, e1))) & ~ (e0 = op(e3, op(e3, e0))) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | (~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | (~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | (~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | (~ (e3 = op(e3, op(e3, e3))) & ~ (e2 = op(e3, op(e3, e2))) & ~ (e1 = op(e3, op(e3, e1))) & ~ (e0 = op(e3, op(e3, e0))) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | (~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | (~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | (~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | (~ (e3 = op(e3, op(e3, e3))) & ~ (e2 = op(e3, op(e3, e2))) & ~ (e1 = op(e3, op(e3, e1))) & ~ (e0 = op(e3, op(e3, e0))) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | (~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | (~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | (~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG159+1.p', ax5)).
fof(f1111, plain, (spl30_110 | spl30_109 | spl30_108 | spl30_107 | spl30_106 | spl30_105 | spl30_104 | spl30_103 | spl30_102 | spl30_101 | spl30_100 | spl30_99 | spl30_98 | spl30_97 | spl30_96 | ~ spl30_4), inference(avatar_split_clause, [], [f398, f446, f952, f962, f972, f982, f992, f1002, f1012, f1022, f1032, f1042, f1052, f1062, f1072, f1082, f1092])).
fof(f398, plain, (~ (e3 = op(e3, e3)) | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f84])).
fof(f1110, plain, (spl30_110 | spl30_109 | spl30_108 | spl30_107 | spl30_106 | spl30_105 | spl30_104 | spl30_103 | spl30_102 | spl30_101 | spl30_100 | spl30_99 | spl30_98 | spl30_97 | spl30_96 | ~ spl30_69), inference(avatar_split_clause, [], [f399, f757, f952, f962, f972, f982, f992, f1002, f1012, f1022, f1032, f1042, f1052, f1062, f1072, f1082, f1092])).
fof(f399, plain, (~ (e0 = op(e3, op(e3, e0))) | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f84])).
fof(f1109, plain, (spl30_110 | spl30_109 | spl30_108 | spl30_107 | spl30_106 | spl30_105 | spl30_104 | spl30_103 | spl30_102 | spl30_101 | spl30_100 | spl30_99 | spl30_98 | spl30_97 | spl30_96 | ~ spl30_68), inference(avatar_split_clause, [], [f400, f752, f952, f962, f972, f982, f992, f1002, f1012, f1022, f1032, f1042, f1052, f1062, f1072, f1082, f1092])).
fof(f400, plain, (~ (e1 = op(e3, op(e3, e1))) | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f84])).
fof(f1108, plain, (spl30_110 | spl30_109 | spl30_108 | spl30_107 | spl30_106 | spl30_105 | spl30_104 | spl30_103 | spl30_102 | spl30_101 | spl30_100 | spl30_99 | spl30_98 | spl30_97 | spl30_96 | ~ spl30_67), inference(avatar_split_clause, [], [f401, f747, f952, f962, f972, f982, f992, f1002, f1012, f1022, f1032, f1042, f1052, f1062, f1072, f1082, f1092])).
fof(f401, plain, (~ (e2 = op(e3, op(e3, e2))) | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f84])).
fof(f1106, plain, (spl30_95 | spl30_94 | spl30_93 | spl30_88 | spl30_87 | spl30_86 | spl30_85 | spl30_80 | spl30_79 | spl30_78 | spl30_77 | spl30_72 | spl30_71 | spl30_70 | spl30_65 | spl30_4), inference(avatar_split_clause, [], [f403, f446, f738, f764, f774, f784, f810, f820, f830, f840, f866, f876, f886, f896, f922, f932, f942])).
fof(f942, plain, (spl30_95 <=> sP15), introduced(avatar_definition, [new_symbols(naming, [spl30_95])])).
fof(f932, plain, (spl30_94 <=> sP16), introduced(avatar_definition, [new_symbols(naming, [spl30_94])])).
fof(f922, plain, (spl30_93 <=> sP17), introduced(avatar_definition, [new_symbols(naming, [spl30_93])])).
fof(f896, plain, (spl30_88 <=> sP18), introduced(avatar_definition, [new_symbols(naming, [spl30_88])])).
fof(f886, plain, (spl30_87 <=> sP19), introduced(avatar_definition, [new_symbols(naming, [spl30_87])])).
fof(f876, plain, (spl30_86 <=> sP20), introduced(avatar_definition, [new_symbols(naming, [spl30_86])])).
fof(f866, plain, (spl30_85 <=> sP21), introduced(avatar_definition, [new_symbols(naming, [spl30_85])])).
fof(f840, plain, (spl30_80 <=> sP22), introduced(avatar_definition, [new_symbols(naming, [spl30_80])])).
fof(f830, plain, (spl30_79 <=> sP23), introduced(avatar_definition, [new_symbols(naming, [spl30_79])])).
fof(f820, plain, (spl30_78 <=> sP24), introduced(avatar_definition, [new_symbols(naming, [spl30_78])])).
fof(f810, plain, (spl30_77 <=> sP25), introduced(avatar_definition, [new_symbols(naming, [spl30_77])])).
fof(f784, plain, (spl30_72 <=> sP26), introduced(avatar_definition, [new_symbols(naming, [spl30_72])])).
fof(f774, plain, (spl30_71 <=> sP27), introduced(avatar_definition, [new_symbols(naming, [spl30_71])])).
fof(f764, plain, (spl30_70 <=> sP28), introduced(avatar_definition, [new_symbols(naming, [spl30_70])])).
fof(f738, plain, (spl30_65 <=> sP29), introduced(avatar_definition, [new_symbols(naming, [spl30_65])])).
fof(f403, plain, ((e3 = op(e3, e3)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15), inference(cnf_transformation, [], [f84])).
fof(f1104, plain, (spl30_95 | spl30_94 | spl30_93 | spl30_88 | spl30_87 | spl30_86 | spl30_85 | spl30_80 | spl30_79 | spl30_78 | spl30_77 | spl30_72 | spl30_71 | spl30_70 | spl30_65 | ~ spl30_69), inference(avatar_split_clause, [], [f405, f757, f738, f764, f774, f784, f810, f820, f830, f840, f866, f876, f886, f896, f922, f932, f942])).
fof(f405, plain, (~ (e0 = op(e3, op(e3, e0))) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15), inference(cnf_transformation, [], [f84])).
fof(f1102, plain, (spl30_95 | spl30_94 | spl30_93 | spl30_88 | spl30_87 | spl30_86 | spl30_85 | spl30_80 | spl30_79 | spl30_78 | spl30_77 | spl30_72 | spl30_71 | spl30_70 | spl30_65 | ~ spl30_67), inference(avatar_split_clause, [], [f407, f747, f738, f764, f774, f784, f810, f820, f830, f840, f866, f876, f886, f896, f922, f932, f942])).
fof(f407, plain, (~ (e2 = op(e3, op(e3, e2))) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15), inference(cnf_transformation, [], [f84])).
fof(f1101, plain, (spl30_95 | spl30_94 | spl30_93 | spl30_88 | spl30_87 | spl30_86 | spl30_85 | spl30_80 | spl30_79 | spl30_78 | spl30_77 | spl30_72 | spl30_71 | spl30_70 | spl30_65 | ~ spl30_66), inference(avatar_split_clause, [], [f408, f742, f738, f764, f774, f784, f810, f820, f830, f840, f866, f876, f886, f896, f922, f932, f942])).
fof(f408, plain, (~ (e3 = op(e3, op(e3, e3))) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15), inference(cnf_transformation, [], [f84])).
fof(f1100, plain, (~ spl30_110 | spl30_61), inference(avatar_split_clause, [], [f391, f689, f1092])).
fof(f391, plain, ((e0 = op(e0, e0)) | ~ sP0), inference(cnf_transformation, [], [f114])).
fof(f114, plain, ((~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP0), inference(nnf_transformation, [], [f54])).
fof(f1099, plain, (~ spl30_110 | ~ spl30_61), inference(avatar_split_clause, [], [f392, f689, f1092])).
fof(f392, plain, (~ (e0 = op(e0, e0)) | ~ sP0), inference(cnf_transformation, [], [f114])).
fof(f1096, plain, (~ spl30_110 | ~ spl30_90), inference(avatar_split_clause, [], [f395, f905, f1092])).
fof(f395, plain, (~ (e2 = op(e0, op(e0, e2))) | ~ sP0), inference(cnf_transformation, [], [f114])).
fof(f1095, plain, (~ spl30_110 | ~ spl30_89), inference(avatar_split_clause, [], [f396, f900, f1092])).
fof(f396, plain, (~ (e3 = op(e0, op(e0, e3))) | ~ sP0), inference(cnf_transformation, [], [f114])).
fof(f1090, plain, (~ spl30_109 | spl30_57), inference(avatar_split_clause, [], [f385, f672, f1082])).
fof(f385, plain, ((e0 = op(e0, e1)) | ~ sP1), inference(cnf_transformation, [], [f113])).
fof(f113, plain, ((~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | ~ sP1), inference(nnf_transformation, [], [f55])).
fof(f1089, plain, (~ spl30_109 | ~ spl30_62), inference(avatar_split_clause, [], [f386, f693, f1082])).
fof(f386, plain, (~ (op(e0, e0) = e1) | ~ sP1), inference(cnf_transformation, [], [f113])).
fof(f1088, plain, (~ spl30_109 | ~ spl30_84), inference(avatar_split_clause, [], [f387, f859, f1082])).
fof(f387, plain, (~ (e0 = op(e1, op(e1, e0))) | ~ sP1), inference(cnf_transformation, [], [f113])).
fof(f1087, plain, (~ spl30_109 | ~ spl30_83), inference(avatar_split_clause, [], [f388, f854, f1082])).
fof(f388, plain, (~ (e1 = op(e1, op(e1, e1))) | ~ sP1), inference(cnf_transformation, [], [f113])).
fof(f1086, plain, (~ spl30_109 | ~ spl30_82), inference(avatar_split_clause, [], [f389, f849, f1082])).
fof(f389, plain, (~ (e2 = op(e1, op(e1, e2))) | ~ sP1), inference(cnf_transformation, [], [f113])).
fof(f1085, plain, (~ spl30_109 | ~ spl30_81), inference(avatar_split_clause, [], [f390, f844, f1082])).
fof(f390, plain, (~ (e3 = op(e1, op(e1, e3))) | ~ sP1), inference(cnf_transformation, [], [f113])).
fof(f1080, plain, (~ spl30_108 | spl30_53), inference(avatar_split_clause, [], [f379, f655, f1072])).
fof(f379, plain, ((e0 = op(e0, e2)) | ~ sP2), inference(cnf_transformation, [], [f112])).
fof(f112, plain, ((~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | ~ sP2), inference(nnf_transformation, [], [f56])).
fof(f1079, plain, (~ spl30_108 | ~ spl30_63), inference(avatar_split_clause, [], [f380, f697, f1072])).
fof(f380, plain, (~ (op(e0, e0) = e2) | ~ sP2), inference(cnf_transformation, [], [f112])).
fof(f1078, plain, (~ spl30_108 | ~ spl30_76), inference(avatar_split_clause, [], [f381, f803, f1072])).
fof(f381, plain, (~ (e0 = op(e2, op(e2, e0))) | ~ sP2), inference(cnf_transformation, [], [f112])).
fof(f1077, plain, (~ spl30_108 | ~ spl30_75), inference(avatar_split_clause, [], [f382, f798, f1072])).
fof(f382, plain, (~ (e1 = op(e2, op(e2, e1))) | ~ sP2), inference(cnf_transformation, [], [f112])).
fof(f1076, plain, (~ spl30_108 | ~ spl30_74), inference(avatar_split_clause, [], [f383, f793, f1072])).
fof(f383, plain, (~ (e2 = op(e2, op(e2, e2))) | ~ sP2), inference(cnf_transformation, [], [f112])).
fof(f1075, plain, (~ spl30_108 | ~ spl30_73), inference(avatar_split_clause, [], [f384, f788, f1072])).
fof(f384, plain, (~ (e3 = op(e2, op(e2, e3))) | ~ sP2), inference(cnf_transformation, [], [f112])).
fof(f1070, plain, (~ spl30_107 | spl30_49), inference(avatar_split_clause, [], [f373, f638, f1062])).
fof(f373, plain, ((e0 = op(e0, e3)) | ~ sP3), inference(cnf_transformation, [], [f111])).
fof(f111, plain, ((~ (e3 = op(e3, op(e3, e3))) & ~ (e2 = op(e3, op(e3, e2))) & ~ (e1 = op(e3, op(e3, e1))) & ~ (e0 = op(e3, op(e3, e0))) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | ~ sP3), inference(nnf_transformation, [], [f57])).
fof(f1069, plain, (~ spl30_107 | ~ spl30_64), inference(avatar_split_clause, [], [f374, f701, f1062])).
fof(f374, plain, (~ (op(e0, e0) = e3) | ~ sP3), inference(cnf_transformation, [], [f111])).
fof(f1068, plain, (~ spl30_107 | ~ spl30_69), inference(avatar_split_clause, [], [f375, f757, f1062])).
fof(f375, plain, (~ (e0 = op(e3, op(e3, e0))) | ~ sP3), inference(cnf_transformation, [], [f111])).
fof(f1067, plain, (~ spl30_107 | ~ spl30_68), inference(avatar_split_clause, [], [f376, f752, f1062])).
fof(f376, plain, (~ (e1 = op(e3, op(e3, e1))) | ~ sP3), inference(cnf_transformation, [], [f111])).
fof(f1066, plain, (~ spl30_107 | ~ spl30_67), inference(avatar_split_clause, [], [f377, f747, f1062])).
fof(f377, plain, (~ (e2 = op(e3, op(e3, e2))) | ~ sP3), inference(cnf_transformation, [], [f111])).
fof(f1065, plain, (~ spl30_107 | ~ spl30_66), inference(avatar_split_clause, [], [f378, f742, f1062])).
fof(f378, plain, (~ (e3 = op(e3, op(e3, e3))) | ~ sP3), inference(cnf_transformation, [], [f111])).
fof(f1060, plain, (~ spl30_106 | spl30_46), inference(avatar_split_clause, [], [f367, f625, f1052])).
fof(f367, plain, ((e1 = op(e1, e0)) | ~ sP4), inference(cnf_transformation, [], [f110])).
fof(f110, plain, ((~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | ~ sP4), inference(nnf_transformation, [], [f58])).
fof(f1059, plain, (~ spl30_106 | ~ spl30_41), inference(avatar_split_clause, [], [f368, f604, f1052])).
fof(f368, plain, (~ (e0 = op(e1, e1)) | ~ sP4), inference(cnf_transformation, [], [f110])).
fof(f1058, plain, (~ spl30_106 | ~ spl30_92), inference(avatar_split_clause, [], [f369, f915, f1052])).
fof(f369, plain, (~ (e0 = op(e0, op(e0, e0))) | ~ sP4), inference(cnf_transformation, [], [f110])).
fof(f1057, plain, (~ spl30_106 | ~ spl30_91), inference(avatar_split_clause, [], [f370, f910, f1052])).
fof(f370, plain, (~ (e1 = op(e0, op(e0, e1))) | ~ sP4), inference(cnf_transformation, [], [f110])).
fof(f1056, plain, (~ spl30_106 | ~ spl30_90), inference(avatar_split_clause, [], [f371, f905, f1052])).
fof(f371, plain, (~ (e2 = op(e0, op(e0, e2))) | ~ sP4), inference(cnf_transformation, [], [f110])).
fof(f1055, plain, (~ spl30_106 | ~ spl30_89), inference(avatar_split_clause, [], [f372, f900, f1052])).
fof(f372, plain, (~ (e3 = op(e0, op(e0, e3))) | ~ sP4), inference(cnf_transformation, [], [f110])).
fof(f1050, plain, (~ spl30_105 | spl30_42), inference(avatar_split_clause, [], [f361, f608, f1042])).
fof(f361, plain, ((e1 = op(e1, e1)) | ~ sP5), inference(cnf_transformation, [], [f109])).
fof(f109, plain, ((~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP5), inference(nnf_transformation, [], [f59])).
fof(f1049, plain, (~ spl30_105 | ~ spl30_42), inference(avatar_split_clause, [], [f362, f608, f1042])).
fof(f362, plain, (~ (e1 = op(e1, e1)) | ~ sP5), inference(cnf_transformation, [], [f109])).
fof(f1046, plain, (~ spl30_105 | ~ spl30_82), inference(avatar_split_clause, [], [f365, f849, f1042])).
fof(f365, plain, (~ (e2 = op(e1, op(e1, e2))) | ~ sP5), inference(cnf_transformation, [], [f109])).
fof(f1040, plain, (~ spl30_104 | spl30_38), inference(avatar_split_clause, [], [f355, f591, f1032])).
fof(f355, plain, ((e1 = op(e1, e2)) | ~ sP6), inference(cnf_transformation, [], [f108])).
fof(f108, plain, ((~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | ~ sP6), inference(nnf_transformation, [], [f60])).
fof(f1039, plain, (~ spl30_104 | ~ spl30_43), inference(avatar_split_clause, [], [f356, f612, f1032])).
fof(f356, plain, (~ (e2 = op(e1, e1)) | ~ sP6), inference(cnf_transformation, [], [f108])).
fof(f1038, plain, (~ spl30_104 | ~ spl30_76), inference(avatar_split_clause, [], [f357, f803, f1032])).
fof(f357, plain, (~ (e0 = op(e2, op(e2, e0))) | ~ sP6), inference(cnf_transformation, [], [f108])).
fof(f1037, plain, (~ spl30_104 | ~ spl30_75), inference(avatar_split_clause, [], [f358, f798, f1032])).
fof(f358, plain, (~ (e1 = op(e2, op(e2, e1))) | ~ sP6), inference(cnf_transformation, [], [f108])).
fof(f1036, plain, (~ spl30_104 | ~ spl30_74), inference(avatar_split_clause, [], [f359, f793, f1032])).
fof(f359, plain, (~ (e2 = op(e2, op(e2, e2))) | ~ sP6), inference(cnf_transformation, [], [f108])).
fof(f1035, plain, (~ spl30_104 | ~ spl30_73), inference(avatar_split_clause, [], [f360, f788, f1032])).
fof(f360, plain, (~ (e3 = op(e2, op(e2, e3))) | ~ sP6), inference(cnf_transformation, [], [f108])).
fof(f1030, plain, (~ spl30_103 | spl30_34), inference(avatar_split_clause, [], [f349, f574, f1022])).
fof(f349, plain, ((e1 = op(e1, e3)) | ~ sP7), inference(cnf_transformation, [], [f107])).
fof(f107, plain, ((~ (e3 = op(e3, op(e3, e3))) & ~ (e2 = op(e3, op(e3, e2))) & ~ (e1 = op(e3, op(e3, e1))) & ~ (e0 = op(e3, op(e3, e0))) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | ~ sP7), inference(nnf_transformation, [], [f61])).
fof(f1029, plain, (~ spl30_103 | ~ spl30_44), inference(avatar_split_clause, [], [f350, f616, f1022])).
fof(f350, plain, (~ (e3 = op(e1, e1)) | ~ sP7), inference(cnf_transformation, [], [f107])).
fof(f1028, plain, (~ spl30_103 | ~ spl30_69), inference(avatar_split_clause, [], [f351, f757, f1022])).
fof(f351, plain, (~ (e0 = op(e3, op(e3, e0))) | ~ sP7), inference(cnf_transformation, [], [f107])).
fof(f1027, plain, (~ spl30_103 | ~ spl30_68), inference(avatar_split_clause, [], [f352, f752, f1022])).
fof(f352, plain, (~ (e1 = op(e3, op(e3, e1))) | ~ sP7), inference(cnf_transformation, [], [f107])).
fof(f1026, plain, (~ spl30_103 | ~ spl30_67), inference(avatar_split_clause, [], [f353, f747, f1022])).
fof(f353, plain, (~ (e2 = op(e3, op(e3, e2))) | ~ sP7), inference(cnf_transformation, [], [f107])).
fof(f1025, plain, (~ spl30_103 | ~ spl30_66), inference(avatar_split_clause, [], [f354, f742, f1022])).
fof(f354, plain, (~ (e3 = op(e3, op(e3, e3))) | ~ sP7), inference(cnf_transformation, [], [f107])).
fof(f1020, plain, (~ spl30_102 | spl30_31), inference(avatar_split_clause, [], [f343, f561, f1012])).
fof(f343, plain, ((e2 = op(e2, e0)) | ~ sP8), inference(cnf_transformation, [], [f106])).
fof(f106, plain, ((~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | ~ sP8), inference(nnf_transformation, [], [f62])).
fof(f1019, plain, (~ spl30_102 | ~ spl30_21), inference(avatar_split_clause, [], [f344, f519, f1012])).
fof(f344, plain, (~ (e0 = op(e2, e2)) | ~ sP8), inference(cnf_transformation, [], [f106])).
fof(f1018, plain, (~ spl30_102 | ~ spl30_92), inference(avatar_split_clause, [], [f345, f915, f1012])).
fof(f345, plain, (~ (e0 = op(e0, op(e0, e0))) | ~ sP8), inference(cnf_transformation, [], [f106])).
fof(f1017, plain, (~ spl30_102 | ~ spl30_91), inference(avatar_split_clause, [], [f346, f910, f1012])).
fof(f346, plain, (~ (e1 = op(e0, op(e0, e1))) | ~ sP8), inference(cnf_transformation, [], [f106])).
fof(f1016, plain, (~ spl30_102 | ~ spl30_90), inference(avatar_split_clause, [], [f347, f905, f1012])).
fof(f347, plain, (~ (e2 = op(e0, op(e0, e2))) | ~ sP8), inference(cnf_transformation, [], [f106])).
fof(f1015, plain, (~ spl30_102 | ~ spl30_89), inference(avatar_split_clause, [], [f348, f900, f1012])).
fof(f348, plain, (~ (e3 = op(e0, op(e0, e3))) | ~ sP8), inference(cnf_transformation, [], [f106])).
fof(f1010, plain, (~ spl30_101 | spl30_27), inference(avatar_split_clause, [], [f337, f544, f1002])).
fof(f337, plain, ((e2 = op(e2, e1)) | ~ sP9), inference(cnf_transformation, [], [f105])).
fof(f105, plain, ((~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | ~ sP9), inference(nnf_transformation, [], [f63])).
fof(f1009, plain, (~ spl30_101 | ~ spl30_22), inference(avatar_split_clause, [], [f338, f523, f1002])).
fof(f338, plain, (~ (e1 = op(e2, e2)) | ~ sP9), inference(cnf_transformation, [], [f105])).
fof(f1008, plain, (~ spl30_101 | ~ spl30_84), inference(avatar_split_clause, [], [f339, f859, f1002])).
fof(f339, plain, (~ (e0 = op(e1, op(e1, e0))) | ~ sP9), inference(cnf_transformation, [], [f105])).
fof(f1007, plain, (~ spl30_101 | ~ spl30_83), inference(avatar_split_clause, [], [f340, f854, f1002])).
fof(f340, plain, (~ (e1 = op(e1, op(e1, e1))) | ~ sP9), inference(cnf_transformation, [], [f105])).
fof(f1006, plain, (~ spl30_101 | ~ spl30_82), inference(avatar_split_clause, [], [f341, f849, f1002])).
fof(f341, plain, (~ (e2 = op(e1, op(e1, e2))) | ~ sP9), inference(cnf_transformation, [], [f105])).
fof(f1005, plain, (~ spl30_101 | ~ spl30_81), inference(avatar_split_clause, [], [f342, f844, f1002])).
fof(f342, plain, (~ (e3 = op(e1, op(e1, e3))) | ~ sP9), inference(cnf_transformation, [], [f105])).
fof(f1000, plain, (~ spl30_100 | spl30_23), inference(avatar_split_clause, [], [f331, f527, f992])).
fof(f331, plain, ((e2 = op(e2, e2)) | ~ sP10), inference(cnf_transformation, [], [f104])).
fof(f104, plain, ((~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP10), inference(nnf_transformation, [], [f64])).
fof(f999, plain, (~ spl30_100 | ~ spl30_23), inference(avatar_split_clause, [], [f332, f527, f992])).
fof(f332, plain, (~ (e2 = op(e2, e2)) | ~ sP10), inference(cnf_transformation, [], [f104])).
fof(f997, plain, (~ spl30_100 | ~ spl30_75), inference(avatar_split_clause, [], [f334, f798, f992])).
fof(f334, plain, (~ (e1 = op(e2, op(e2, e1))) | ~ sP10), inference(cnf_transformation, [], [f104])).
fof(f995, plain, (~ spl30_100 | ~ spl30_73), inference(avatar_split_clause, [], [f336, f788, f992])).
fof(f336, plain, (~ (e3 = op(e2, op(e2, e3))) | ~ sP10), inference(cnf_transformation, [], [f104])).
fof(f990, plain, (~ spl30_99 | spl30_19), inference(avatar_split_clause, [], [f325, f510, f982])).
fof(f325, plain, ((e2 = op(e2, e3)) | ~ sP11), inference(cnf_transformation, [], [f103])).
fof(f103, plain, ((~ (e3 = op(e3, op(e3, e3))) & ~ (e2 = op(e3, op(e3, e2))) & ~ (e1 = op(e3, op(e3, e1))) & ~ (e0 = op(e3, op(e3, e0))) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | ~ sP11), inference(nnf_transformation, [], [f65])).
fof(f989, plain, (~ spl30_99 | ~ spl30_24), inference(avatar_split_clause, [], [f326, f531, f982])).
fof(f326, plain, (~ (e3 = op(e2, e2)) | ~ sP11), inference(cnf_transformation, [], [f103])).
fof(f988, plain, (~ spl30_99 | ~ spl30_69), inference(avatar_split_clause, [], [f327, f757, f982])).
fof(f327, plain, (~ (e0 = op(e3, op(e3, e0))) | ~ sP11), inference(cnf_transformation, [], [f103])).
fof(f987, plain, (~ spl30_99 | ~ spl30_68), inference(avatar_split_clause, [], [f328, f752, f982])).
fof(f328, plain, (~ (e1 = op(e3, op(e3, e1))) | ~ sP11), inference(cnf_transformation, [], [f103])).
fof(f986, plain, (~ spl30_99 | ~ spl30_67), inference(avatar_split_clause, [], [f329, f747, f982])).
fof(f329, plain, (~ (e2 = op(e3, op(e3, e2))) | ~ sP11), inference(cnf_transformation, [], [f103])).
fof(f985, plain, (~ spl30_99 | ~ spl30_66), inference(avatar_split_clause, [], [f330, f742, f982])).
fof(f330, plain, (~ (e3 = op(e3, op(e3, e3))) | ~ sP11), inference(cnf_transformation, [], [f103])).
fof(f980, plain, (~ spl30_98 | spl30_16), inference(avatar_split_clause, [], [f319, f497, f972])).
fof(f319, plain, ((e3 = op(e3, e0)) | ~ sP12), inference(cnf_transformation, [], [f102])).
fof(f102, plain, ((~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | ~ sP12), inference(nnf_transformation, [], [f66])).
fof(f979, plain, (~ spl30_98 | ~ spl30_1), inference(avatar_split_clause, [], [f320, f434, f972])).
fof(f320, plain, (~ (e0 = op(e3, e3)) | ~ sP12), inference(cnf_transformation, [], [f102])).
fof(f978, plain, (~ spl30_98 | ~ spl30_92), inference(avatar_split_clause, [], [f321, f915, f972])).
fof(f321, plain, (~ (e0 = op(e0, op(e0, e0))) | ~ sP12), inference(cnf_transformation, [], [f102])).
fof(f977, plain, (~ spl30_98 | ~ spl30_91), inference(avatar_split_clause, [], [f322, f910, f972])).
fof(f322, plain, (~ (e1 = op(e0, op(e0, e1))) | ~ sP12), inference(cnf_transformation, [], [f102])).
fof(f976, plain, (~ spl30_98 | ~ spl30_90), inference(avatar_split_clause, [], [f323, f905, f972])).
fof(f323, plain, (~ (e2 = op(e0, op(e0, e2))) | ~ sP12), inference(cnf_transformation, [], [f102])).
fof(f975, plain, (~ spl30_98 | ~ spl30_89), inference(avatar_split_clause, [], [f324, f900, f972])).
fof(f324, plain, (~ (e3 = op(e0, op(e0, e3))) | ~ sP12), inference(cnf_transformation, [], [f102])).
fof(f970, plain, (~ spl30_97 | spl30_12), inference(avatar_split_clause, [], [f313, f480, f962])).
fof(f313, plain, ((e3 = op(e3, e1)) | ~ sP13), inference(cnf_transformation, [], [f101])).
fof(f101, plain, ((~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | ~ sP13), inference(nnf_transformation, [], [f67])).
fof(f969, plain, (~ spl30_97 | ~ spl30_2), inference(avatar_split_clause, [], [f314, f438, f962])).
fof(f314, plain, (~ (e1 = op(e3, e3)) | ~ sP13), inference(cnf_transformation, [], [f101])).
fof(f968, plain, (~ spl30_97 | ~ spl30_84), inference(avatar_split_clause, [], [f315, f859, f962])).
fof(f315, plain, (~ (e0 = op(e1, op(e1, e0))) | ~ sP13), inference(cnf_transformation, [], [f101])).
fof(f967, plain, (~ spl30_97 | ~ spl30_83), inference(avatar_split_clause, [], [f316, f854, f962])).
fof(f316, plain, (~ (e1 = op(e1, op(e1, e1))) | ~ sP13), inference(cnf_transformation, [], [f101])).
fof(f966, plain, (~ spl30_97 | ~ spl30_82), inference(avatar_split_clause, [], [f317, f849, f962])).
fof(f317, plain, (~ (e2 = op(e1, op(e1, e2))) | ~ sP13), inference(cnf_transformation, [], [f101])).
fof(f965, plain, (~ spl30_97 | ~ spl30_81), inference(avatar_split_clause, [], [f318, f844, f962])).
fof(f318, plain, (~ (e3 = op(e1, op(e1, e3))) | ~ sP13), inference(cnf_transformation, [], [f101])).
fof(f960, plain, (~ spl30_96 | spl30_8), inference(avatar_split_clause, [], [f307, f463, f952])).
fof(f307, plain, ((e3 = op(e3, e2)) | ~ sP14), inference(cnf_transformation, [], [f100])).
fof(f100, plain, ((~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | ~ sP14), inference(nnf_transformation, [], [f68])).
fof(f959, plain, (~ spl30_96 | ~ spl30_3), inference(avatar_split_clause, [], [f308, f442, f952])).
fof(f308, plain, (~ (e2 = op(e3, e3)) | ~ sP14), inference(cnf_transformation, [], [f100])).
fof(f958, plain, (~ spl30_96 | ~ spl30_76), inference(avatar_split_clause, [], [f309, f803, f952])).
fof(f309, plain, (~ (e0 = op(e2, op(e2, e0))) | ~ sP14), inference(cnf_transformation, [], [f100])).
fof(f957, plain, (~ spl30_96 | ~ spl30_75), inference(avatar_split_clause, [], [f310, f798, f952])).
fof(f310, plain, (~ (e1 = op(e2, op(e2, e1))) | ~ sP14), inference(cnf_transformation, [], [f100])).
fof(f956, plain, (~ spl30_96 | ~ spl30_74), inference(avatar_split_clause, [], [f311, f793, f952])).
fof(f311, plain, (~ (e2 = op(e2, op(e2, e2))) | ~ sP14), inference(cnf_transformation, [], [f100])).
fof(f955, plain, (~ spl30_96 | ~ spl30_73), inference(avatar_split_clause, [], [f312, f788, f952])).
fof(f312, plain, (~ (e3 = op(e2, op(e2, e3))) | ~ sP14), inference(cnf_transformation, [], [f100])).
fof(f950, plain, (~ spl30_95 | spl30_61), inference(avatar_split_clause, [], [f301, f689, f942])).
fof(f301, plain, ((e0 = op(e0, e0)) | ~ sP15), inference(cnf_transformation, [], [f99])).
fof(f99, plain, ((~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP15), inference(nnf_transformation, [], [f69])).
fof(f948, plain, (~ spl30_95 | ~ spl30_92), inference(avatar_split_clause, [], [f303, f915, f942])).
fof(f303, plain, (~ (e0 = op(e0, op(e0, e0))) | ~ sP15), inference(cnf_transformation, [], [f99])).
fof(f940, plain, (~ spl30_94 | spl30_57), inference(avatar_split_clause, [], [f295, f672, f932])).
fof(f295, plain, ((e0 = op(e0, e1)) | ~ sP16), inference(cnf_transformation, [], [f98])).
fof(f98, plain, ((~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & (e0 = op(e1, e0)) & (e0 = op(e0, e1))) | ~ sP16), inference(nnf_transformation, [], [f70])).
fof(f939, plain, (~ spl30_94 | spl30_45), inference(avatar_split_clause, [], [f296, f621, f932])).
fof(f296, plain, ((e0 = op(e1, e0)) | ~ sP16), inference(cnf_transformation, [], [f98])).
fof(f937, plain, (~ spl30_94 | ~ spl30_91), inference(avatar_split_clause, [], [f298, f910, f932])).
fof(f298, plain, (~ (e1 = op(e0, op(e0, e1))) | ~ sP16), inference(cnf_transformation, [], [f98])).
fof(f936, plain, (~ spl30_94 | ~ spl30_90), inference(avatar_split_clause, [], [f299, f905, f932])).
fof(f299, plain, (~ (e2 = op(e0, op(e0, e2))) | ~ sP16), inference(cnf_transformation, [], [f98])).
fof(f935, plain, (~ spl30_94 | ~ spl30_89), inference(avatar_split_clause, [], [f300, f900, f932])).
fof(f300, plain, (~ (e3 = op(e0, op(e0, e3))) | ~ sP16), inference(cnf_transformation, [], [f98])).
fof(f930, plain, (~ spl30_93 | spl30_53), inference(avatar_split_clause, [], [f289, f655, f922])).
fof(f289, plain, ((e0 = op(e0, e2)) | ~ sP17), inference(cnf_transformation, [], [f97])).
fof(f97, plain, ((~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & (e0 = op(e2, e0)) & (e0 = op(e0, e2))) | ~ sP17), inference(nnf_transformation, [], [f71])).
fof(f929, plain, (~ spl30_93 | spl30_29), inference(avatar_split_clause, [], [f290, f553, f922])).
fof(f290, plain, ((e0 = op(e2, e0)) | ~ sP17), inference(cnf_transformation, [], [f97])).
fof(f928, plain, (~ spl30_93 | ~ spl30_92), inference(avatar_split_clause, [], [f291, f915, f922])).
fof(f291, plain, (~ (e0 = op(e0, op(e0, e0))) | ~ sP17), inference(cnf_transformation, [], [f97])).
fof(f927, plain, (~ spl30_93 | ~ spl30_91), inference(avatar_split_clause, [], [f292, f910, f922])).
fof(f292, plain, (~ (e1 = op(e0, op(e0, e1))) | ~ sP17), inference(cnf_transformation, [], [f97])).
fof(f926, plain, (~ spl30_93 | ~ spl30_90), inference(avatar_split_clause, [], [f293, f905, f922])).
fof(f293, plain, (~ (e2 = op(e0, op(e0, e2))) | ~ sP17), inference(cnf_transformation, [], [f97])).
fof(f925, plain, (~ spl30_93 | ~ spl30_89), inference(avatar_split_clause, [], [f294, f900, f922])).
fof(f294, plain, (~ (e3 = op(e0, op(e0, e3))) | ~ sP17), inference(cnf_transformation, [], [f97])).
fof(f920, plain, (~ spl30_88 | spl30_49), inference(avatar_split_clause, [], [f283, f638, f896])).
fof(f283, plain, ((e0 = op(e0, e3)) | ~ sP18), inference(cnf_transformation, [], [f96])).
fof(f96, plain, ((~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & (e0 = op(e3, e0)) & (e0 = op(e0, e3))) | ~ sP18), inference(nnf_transformation, [], [f72])).
fof(f919, plain, (~ spl30_88 | spl30_13), inference(avatar_split_clause, [], [f284, f485, f896])).
fof(f284, plain, ((e0 = op(e3, e0)) | ~ sP18), inference(cnf_transformation, [], [f96])).
fof(f918, plain, (~ spl30_88 | ~ spl30_92), inference(avatar_split_clause, [], [f285, f915, f896])).
fof(f285, plain, (~ (e0 = op(e0, op(e0, e0))) | ~ sP18), inference(cnf_transformation, [], [f96])).
fof(f913, plain, (~ spl30_88 | ~ spl30_91), inference(avatar_split_clause, [], [f286, f910, f896])).
fof(f286, plain, (~ (e1 = op(e0, op(e0, e1))) | ~ sP18), inference(cnf_transformation, [], [f96])).
fof(f908, plain, (~ spl30_88 | ~ spl30_90), inference(avatar_split_clause, [], [f287, f905, f896])).
fof(f287, plain, (~ (e2 = op(e0, op(e0, e2))) | ~ sP18), inference(cnf_transformation, [], [f96])).
fof(f903, plain, (~ spl30_88 | ~ spl30_89), inference(avatar_split_clause, [], [f288, f900, f896])).
fof(f288, plain, (~ (e3 = op(e0, op(e0, e3))) | ~ sP18), inference(cnf_transformation, [], [f96])).
fof(f894, plain, (~ spl30_87 | spl30_46), inference(avatar_split_clause, [], [f277, f625, f886])).
fof(f277, plain, ((e1 = op(e1, e0)) | ~ sP19), inference(cnf_transformation, [], [f95])).
fof(f95, plain, ((~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & (e1 = op(e0, e1)) & (e1 = op(e1, e0))) | ~ sP19), inference(nnf_transformation, [], [f73])).
fof(f893, plain, (~ spl30_87 | spl30_58), inference(avatar_split_clause, [], [f278, f676, f886])).
fof(f278, plain, ((e1 = op(e0, e1)) | ~ sP19), inference(cnf_transformation, [], [f95])).
fof(f892, plain, (~ spl30_87 | ~ spl30_84), inference(avatar_split_clause, [], [f279, f859, f886])).
fof(f279, plain, (~ (e0 = op(e1, op(e1, e0))) | ~ sP19), inference(cnf_transformation, [], [f95])).
fof(f891, plain, (~ spl30_87 | ~ spl30_83), inference(avatar_split_clause, [], [f280, f854, f886])).
fof(f280, plain, (~ (e1 = op(e1, op(e1, e1))) | ~ sP19), inference(cnf_transformation, [], [f95])).
fof(f890, plain, (~ spl30_87 | ~ spl30_82), inference(avatar_split_clause, [], [f281, f849, f886])).
fof(f281, plain, (~ (e2 = op(e1, op(e1, e2))) | ~ sP19), inference(cnf_transformation, [], [f95])).
fof(f889, plain, (~ spl30_87 | ~ spl30_81), inference(avatar_split_clause, [], [f282, f844, f886])).
fof(f282, plain, (~ (e3 = op(e1, op(e1, e3))) | ~ sP19), inference(cnf_transformation, [], [f95])).
fof(f884, plain, (~ spl30_86 | spl30_42), inference(avatar_split_clause, [], [f271, f608, f876])).
fof(f271, plain, ((e1 = op(e1, e1)) | ~ sP20), inference(cnf_transformation, [], [f94])).
fof(f94, plain, ((~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP20), inference(nnf_transformation, [], [f74])).
fof(f881, plain, (~ spl30_86 | ~ spl30_83), inference(avatar_split_clause, [], [f274, f854, f876])).
fof(f274, plain, (~ (e1 = op(e1, op(e1, e1))) | ~ sP20), inference(cnf_transformation, [], [f94])).
fof(f874, plain, (~ spl30_85 | spl30_38), inference(avatar_split_clause, [], [f265, f591, f866])).
fof(f265, plain, ((e1 = op(e1, e2)) | ~ sP21), inference(cnf_transformation, [], [f93])).
fof(f93, plain, ((~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & (e1 = op(e2, e1)) & (e1 = op(e1, e2))) | ~ sP21), inference(nnf_transformation, [], [f75])).
fof(f873, plain, (~ spl30_85 | spl30_26), inference(avatar_split_clause, [], [f266, f540, f866])).
fof(f266, plain, ((e1 = op(e2, e1)) | ~ sP21), inference(cnf_transformation, [], [f93])).
fof(f872, plain, (~ spl30_85 | ~ spl30_84), inference(avatar_split_clause, [], [f267, f859, f866])).
fof(f267, plain, (~ (e0 = op(e1, op(e1, e0))) | ~ sP21), inference(cnf_transformation, [], [f93])).
fof(f871, plain, (~ spl30_85 | ~ spl30_83), inference(avatar_split_clause, [], [f268, f854, f866])).
fof(f268, plain, (~ (e1 = op(e1, op(e1, e1))) | ~ sP21), inference(cnf_transformation, [], [f93])).
fof(f870, plain, (~ spl30_85 | ~ spl30_82), inference(avatar_split_clause, [], [f269, f849, f866])).
fof(f269, plain, (~ (e2 = op(e1, op(e1, e2))) | ~ sP21), inference(cnf_transformation, [], [f93])).
fof(f869, plain, (~ spl30_85 | ~ spl30_81), inference(avatar_split_clause, [], [f270, f844, f866])).
fof(f270, plain, (~ (e3 = op(e1, op(e1, e3))) | ~ sP21), inference(cnf_transformation, [], [f93])).
fof(f864, plain, (~ spl30_80 | spl30_34), inference(avatar_split_clause, [], [f259, f574, f840])).
fof(f259, plain, ((e1 = op(e1, e3)) | ~ sP22), inference(cnf_transformation, [], [f92])).
fof(f92, plain, ((~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & (e1 = op(e3, e1)) & (e1 = op(e1, e3))) | ~ sP22), inference(nnf_transformation, [], [f76])).
fof(f863, plain, (~ spl30_80 | spl30_10), inference(avatar_split_clause, [], [f260, f472, f840])).
fof(f260, plain, ((e1 = op(e3, e1)) | ~ sP22), inference(cnf_transformation, [], [f92])).
fof(f862, plain, (~ spl30_80 | ~ spl30_84), inference(avatar_split_clause, [], [f261, f859, f840])).
fof(f261, plain, (~ (e0 = op(e1, op(e1, e0))) | ~ sP22), inference(cnf_transformation, [], [f92])).
fof(f857, plain, (~ spl30_80 | ~ spl30_83), inference(avatar_split_clause, [], [f262, f854, f840])).
fof(f262, plain, (~ (e1 = op(e1, op(e1, e1))) | ~ sP22), inference(cnf_transformation, [], [f92])).
fof(f852, plain, (~ spl30_80 | ~ spl30_82), inference(avatar_split_clause, [], [f263, f849, f840])).
fof(f263, plain, (~ (e2 = op(e1, op(e1, e2))) | ~ sP22), inference(cnf_transformation, [], [f92])).
fof(f847, plain, (~ spl30_80 | ~ spl30_81), inference(avatar_split_clause, [], [f264, f844, f840])).
fof(f264, plain, (~ (e3 = op(e1, op(e1, e3))) | ~ sP22), inference(cnf_transformation, [], [f92])).
fof(f838, plain, (~ spl30_79 | spl30_31), inference(avatar_split_clause, [], [f253, f561, f830])).
fof(f253, plain, ((e2 = op(e2, e0)) | ~ sP23), inference(cnf_transformation, [], [f91])).
fof(f91, plain, ((~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & (e2 = op(e0, e2)) & (e2 = op(e2, e0))) | ~ sP23), inference(nnf_transformation, [], [f77])).
fof(f837, plain, (~ spl30_79 | spl30_55), inference(avatar_split_clause, [], [f254, f663, f830])).
fof(f254, plain, ((e2 = op(e0, e2)) | ~ sP23), inference(cnf_transformation, [], [f91])).
fof(f836, plain, (~ spl30_79 | ~ spl30_76), inference(avatar_split_clause, [], [f255, f803, f830])).
fof(f255, plain, (~ (e0 = op(e2, op(e2, e0))) | ~ sP23), inference(cnf_transformation, [], [f91])).
fof(f835, plain, (~ spl30_79 | ~ spl30_75), inference(avatar_split_clause, [], [f256, f798, f830])).
fof(f256, plain, (~ (e1 = op(e2, op(e2, e1))) | ~ sP23), inference(cnf_transformation, [], [f91])).
fof(f834, plain, (~ spl30_79 | ~ spl30_74), inference(avatar_split_clause, [], [f257, f793, f830])).
fof(f257, plain, (~ (e2 = op(e2, op(e2, e2))) | ~ sP23), inference(cnf_transformation, [], [f91])).
fof(f833, plain, (~ spl30_79 | ~ spl30_73), inference(avatar_split_clause, [], [f258, f788, f830])).
fof(f258, plain, (~ (e3 = op(e2, op(e2, e3))) | ~ sP23), inference(cnf_transformation, [], [f91])).
fof(f828, plain, (~ spl30_78 | spl30_27), inference(avatar_split_clause, [], [f247, f544, f820])).
fof(f247, plain, ((e2 = op(e2, e1)) | ~ sP24), inference(cnf_transformation, [], [f90])).
fof(f90, plain, ((~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & (e2 = op(e1, e2)) & (e2 = op(e2, e1))) | ~ sP24), inference(nnf_transformation, [], [f78])).
fof(f827, plain, (~ spl30_78 | spl30_39), inference(avatar_split_clause, [], [f248, f595, f820])).
fof(f248, plain, ((e2 = op(e1, e2)) | ~ sP24), inference(cnf_transformation, [], [f90])).
fof(f826, plain, (~ spl30_78 | ~ spl30_76), inference(avatar_split_clause, [], [f249, f803, f820])).
fof(f249, plain, (~ (e0 = op(e2, op(e2, e0))) | ~ sP24), inference(cnf_transformation, [], [f90])).
fof(f825, plain, (~ spl30_78 | ~ spl30_75), inference(avatar_split_clause, [], [f250, f798, f820])).
fof(f250, plain, (~ (e1 = op(e2, op(e2, e1))) | ~ sP24), inference(cnf_transformation, [], [f90])).
fof(f823, plain, (~ spl30_78 | ~ spl30_73), inference(avatar_split_clause, [], [f252, f788, f820])).
fof(f252, plain, (~ (e3 = op(e2, op(e2, e3))) | ~ sP24), inference(cnf_transformation, [], [f90])).
fof(f818, plain, (~ spl30_77 | spl30_23), inference(avatar_split_clause, [], [f241, f527, f810])).
fof(f241, plain, ((e2 = op(e2, e2)) | ~ sP25), inference(cnf_transformation, [], [f89])).
fof(f89, plain, ((~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP25), inference(nnf_transformation, [], [f79])).
fof(f814, plain, (~ spl30_77 | ~ spl30_74), inference(avatar_split_clause, [], [f245, f793, f810])).
fof(f245, plain, (~ (e2 = op(e2, op(e2, e2))) | ~ sP25), inference(cnf_transformation, [], [f89])).
fof(f808, plain, (~ spl30_72 | spl30_19), inference(avatar_split_clause, [], [f235, f510, f784])).
fof(f235, plain, ((e2 = op(e2, e3)) | ~ sP26), inference(cnf_transformation, [], [f88])).
fof(f88, plain, ((~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & (e2 = op(e3, e2)) & (e2 = op(e2, e3))) | ~ sP26), inference(nnf_transformation, [], [f80])).
fof(f807, plain, (~ spl30_72 | spl30_7), inference(avatar_split_clause, [], [f236, f459, f784])).
fof(f236, plain, ((e2 = op(e3, e2)) | ~ sP26), inference(cnf_transformation, [], [f88])).
fof(f806, plain, (~ spl30_72 | ~ spl30_76), inference(avatar_split_clause, [], [f237, f803, f784])).
fof(f237, plain, (~ (e0 = op(e2, op(e2, e0))) | ~ sP26), inference(cnf_transformation, [], [f88])).
fof(f801, plain, (~ spl30_72 | ~ spl30_75), inference(avatar_split_clause, [], [f238, f798, f784])).
fof(f238, plain, (~ (e1 = op(e2, op(e2, e1))) | ~ sP26), inference(cnf_transformation, [], [f88])).
fof(f796, plain, (~ spl30_72 | ~ spl30_74), inference(avatar_split_clause, [], [f239, f793, f784])).
fof(f239, plain, (~ (e2 = op(e2, op(e2, e2))) | ~ sP26), inference(cnf_transformation, [], [f88])).
fof(f791, plain, (~ spl30_72 | ~ spl30_73), inference(avatar_split_clause, [], [f240, f788, f784])).
fof(f240, plain, (~ (e3 = op(e2, op(e2, e3))) | ~ sP26), inference(cnf_transformation, [], [f88])).
fof(f782, plain, (~ spl30_71 | spl30_16), inference(avatar_split_clause, [], [f229, f497, f774])).
fof(f229, plain, ((e3 = op(e3, e0)) | ~ sP27), inference(cnf_transformation, [], [f87])).
fof(f87, plain, ((~ (e3 = op(e3, op(e3, e3))) & ~ (e2 = op(e3, op(e3, e2))) & ~ (e1 = op(e3, op(e3, e1))) & ~ (e0 = op(e3, op(e3, e0))) & (e3 = op(e0, e3)) & (e3 = op(e3, e0))) | ~ sP27), inference(nnf_transformation, [], [f81])).
fof(f781, plain, (~ spl30_71 | spl30_52), inference(avatar_split_clause, [], [f230, f650, f774])).
fof(f230, plain, ((e3 = op(e0, e3)) | ~ sP27), inference(cnf_transformation, [], [f87])).
fof(f780, plain, (~ spl30_71 | ~ spl30_69), inference(avatar_split_clause, [], [f231, f757, f774])).
fof(f231, plain, (~ (e0 = op(e3, op(e3, e0))) | ~ sP27), inference(cnf_transformation, [], [f87])).
fof(f779, plain, (~ spl30_71 | ~ spl30_68), inference(avatar_split_clause, [], [f232, f752, f774])).
fof(f232, plain, (~ (e1 = op(e3, op(e3, e1))) | ~ sP27), inference(cnf_transformation, [], [f87])).
fof(f778, plain, (~ spl30_71 | ~ spl30_67), inference(avatar_split_clause, [], [f233, f747, f774])).
fof(f233, plain, (~ (e2 = op(e3, op(e3, e2))) | ~ sP27), inference(cnf_transformation, [], [f87])).
fof(f772, plain, (~ spl30_70 | spl30_12), inference(avatar_split_clause, [], [f223, f480, f764])).
fof(f223, plain, ((e3 = op(e3, e1)) | ~ sP28), inference(cnf_transformation, [], [f86])).
fof(f86, plain, ((~ (e3 = op(e3, op(e3, e3))) & ~ (e2 = op(e3, op(e3, e2))) & ~ (e1 = op(e3, op(e3, e1))) & ~ (e0 = op(e3, op(e3, e0))) & (e3 = op(e1, e3)) & (e3 = op(e3, e1))) | ~ sP28), inference(nnf_transformation, [], [f82])).
fof(f771, plain, (~ spl30_70 | spl30_36), inference(avatar_split_clause, [], [f224, f582, f764])).
fof(f224, plain, ((e3 = op(e1, e3)) | ~ sP28), inference(cnf_transformation, [], [f86])).
fof(f770, plain, (~ spl30_70 | ~ spl30_69), inference(avatar_split_clause, [], [f225, f757, f764])).
fof(f225, plain, (~ (e0 = op(e3, op(e3, e0))) | ~ sP28), inference(cnf_transformation, [], [f86])).
fof(f769, plain, (~ spl30_70 | ~ spl30_68), inference(avatar_split_clause, [], [f226, f752, f764])).
fof(f226, plain, (~ (e1 = op(e3, op(e3, e1))) | ~ sP28), inference(cnf_transformation, [], [f86])).
fof(f768, plain, (~ spl30_70 | ~ spl30_67), inference(avatar_split_clause, [], [f227, f747, f764])).
fof(f227, plain, (~ (e2 = op(e3, op(e3, e2))) | ~ sP28), inference(cnf_transformation, [], [f86])).
fof(f767, plain, (~ spl30_70 | ~ spl30_66), inference(avatar_split_clause, [], [f228, f742, f764])).
fof(f228, plain, (~ (e3 = op(e3, op(e3, e3))) | ~ sP28), inference(cnf_transformation, [], [f86])).
fof(f762, plain, (~ spl30_65 | spl30_8), inference(avatar_split_clause, [], [f217, f463, f738])).
fof(f217, plain, ((e3 = op(e3, e2)) | ~ sP29), inference(cnf_transformation, [], [f85])).
fof(f85, plain, ((~ (e3 = op(e3, op(e3, e3))) & ~ (e2 = op(e3, op(e3, e2))) & ~ (e1 = op(e3, op(e3, e1))) & ~ (e0 = op(e3, op(e3, e0))) & (e3 = op(e2, e3)) & (e3 = op(e3, e2))) | ~ sP29), inference(nnf_transformation, [], [f83])).
fof(f761, plain, (~ spl30_65 | spl30_20), inference(avatar_split_clause, [], [f218, f514, f738])).
fof(f218, plain, ((e3 = op(e2, e3)) | ~ sP29), inference(cnf_transformation, [], [f85])).
fof(f760, plain, (~ spl30_65 | ~ spl30_69), inference(avatar_split_clause, [], [f219, f757, f738])).
fof(f219, plain, (~ (e0 = op(e3, op(e3, e0))) | ~ sP29), inference(cnf_transformation, [], [f85])).
fof(f755, plain, (~ spl30_65 | ~ spl30_68), inference(avatar_split_clause, [], [f220, f752, f738])).
fof(f220, plain, (~ (e1 = op(e3, op(e3, e1))) | ~ sP29), inference(cnf_transformation, [], [f85])).
fof(f750, plain, (~ spl30_65 | ~ spl30_67), inference(avatar_split_clause, [], [f221, f747, f738])).
fof(f221, plain, (~ (e2 = op(e3, op(e3, e2))) | ~ sP29), inference(cnf_transformation, [], [f85])).
fof(f745, plain, (~ spl30_65 | ~ spl30_66), inference(avatar_split_clause, [], [f222, f742, f738])).
fof(f222, plain, (~ (e3 = op(e3, op(e3, e3))) | ~ sP29), inference(cnf_transformation, [], [f85])).
fof(f736, plain, (spl30_61 | spl30_57 | spl30_53 | spl30_49), inference(avatar_split_clause, [], [f131, f638, f655, f672, f689])).
fof(f131, plain, ((e0 = op(e0, e3)) | (e0 = op(e0, e2)) | (e0 = op(e0, e1)) | (e0 = op(e0, e0))), inference(cnf_transformation, [], [f2])).
fof(f2, plain, (((e3 = op(e3, e3)) | (e3 = op(e2, e3)) | (e3 = op(e1, e3)) | (e3 = op(e0, e3))) & ((e3 = op(e3, e3)) | (e3 = op(e3, e2)) | (e3 = op(e3, e1)) | (e3 = op(e3, e0))) & ((e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(e0, e3))) & ((e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, e0))) & ((e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(e0, e3))) & ((e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, e0))) & ((e0 = op(e3, e3)) | (e0 = op(e2, e3)) | (e0 = op(e1, e3)) | (e0 = op(e0, e3))) & ((e0 = op(e3, e3)) | (e0 = op(e3, e2)) | (e0 = op(e3, e1)) | (e0 = op(e3, e0))) & ((e3 = op(e3, e2)) | (e3 = op(e2, e2)) | (e3 = op(e1, e2)) | (e3 = op(e0, e2))) & ((e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, e0))) & ((e2 = op(e3, e2)) | (e2 = op(e2, e2)) | (e2 = op(e1, e2)) | (e2 = op(e0, e2))) & ((e2 = op(e2, e3)) | (e2 = op(e2, e2)) | (e2 = op(e2, e1)) | (e2 = op(e2, e0))) & ((e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(e0, e2))) & ((e1 = op(e2, e3)) | (e1 = op(e2, e2)) | (e1 = op(e2, e1)) | (e1 = op(e2, e0))) & ((e0 = op(e3, e2)) | (e0 = op(e2, e2)) | (e0 = op(e1, e2)) | (e0 = op(e0, e2))) & ((e0 = op(e2, e3)) | (e0 = op(e2, e2)) | (e0 = op(e2, e1)) | (e0 = op(e2, e0))) & ((e3 = op(e3, e1)) | (e3 = op(e2, e1)) | (e3 = op(e1, e1)) | (e3 = op(e0, e1))) & ((e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, e0))) & ((e2 = op(e3, e1)) | (e2 = op(e2, e1)) | (e2 = op(e1, e1)) | (e2 = op(e0, e1))) & ((e2 = op(e1, e3)) | (e2 = op(e1, e2)) | (e2 = op(e1, e1)) | (e2 = op(e1, e0))) & ((e1 = op(e3, e1)) | (e1 = op(e2, e1)) | (e1 = op(e1, e1)) | (e1 = op(e0, e1))) & ((e1 = op(e1, e3)) | (e1 = op(e1, e2)) | (e1 = op(e1, e1)) | (e1 = op(e1, e0))) & ((e0 = op(e3, e1)) | (e0 = op(e2, e1)) | (e0 = op(e1, e1)) | (e0 = op(e0, e1))) & ((e0 = op(e1, e3)) | (e0 = op(e1, e2)) | (e0 = op(e1, e1)) | (e0 = op(e1, e0))) & ((e3 = op(e3, e0)) | (e3 = op(e2, e0)) | (e3 = op(e1, e0)) | (op(e0, e0) = e3)) & ((e3 = op(e0, e3)) | (e3 = op(e0, e2)) | (e3 = op(e0, e1)) | (op(e0, e0) = e3)) & ((e2 = op(e3, e0)) | (e2 = op(e2, e0)) | (e2 = op(e1, e0)) | (op(e0, e0) = e2)) & ((e2 = op(e0, e3)) | (e2 = op(e0, e2)) | (e2 = op(e0, e1)) | (op(e0, e0) = e2)) & ((e1 = op(e3, e0)) | (e1 = op(e2, e0)) | (e1 = op(e1, e0)) | (op(e0, e0) = e1)) & ((e1 = op(e0, e3)) | (e1 = op(e0, e2)) | (e1 = op(e0, e1)) | (op(e0, e0) = e1)) & ((e0 = op(e3, e0)) | (e0 = op(e2, e0)) | (e0 = op(e1, e0)) | (e0 = op(e0, e0))) & ((e0 = op(e0, e3)) | (e0 = op(e0, e2)) | (e0 = op(e0, e1)) | (e0 = op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG159+1.p', ax2)).
fof(f735, plain, (spl30_61 | spl30_45 | spl30_29 | spl30_13), inference(avatar_split_clause, [], [f132, f485, f553, f621, f689])).
fof(f132, plain, ((e0 = op(e3, e0)) | (e0 = op(e2, e0)) | (e0 = op(e1, e0)) | (e0 = op(e0, e0))), inference(cnf_transformation, [], [f2])).
fof(f734, plain, (spl30_62 | spl30_58 | spl30_54 | spl30_50), inference(avatar_split_clause, [], [f133, f642, f659, f676, f693])).
fof(f133, plain, ((e1 = op(e0, e3)) | (e1 = op(e0, e2)) | (e1 = op(e0, e1)) | (op(e0, e0) = e1)), inference(cnf_transformation, [], [f2])).
fof(f733, plain, (spl30_62 | spl30_46 | spl30_30 | spl30_14), inference(avatar_split_clause, [], [f134, f489, f557, f625, f693])).
fof(f134, plain, ((e1 = op(e3, e0)) | (e1 = op(e2, e0)) | (e1 = op(e1, e0)) | (op(e0, e0) = e1)), inference(cnf_transformation, [], [f2])).
fof(f732, plain, (spl30_63 | spl30_59 | spl30_55 | spl30_51), inference(avatar_split_clause, [], [f135, f646, f663, f680, f697])).
fof(f135, plain, ((e2 = op(e0, e3)) | (e2 = op(e0, e2)) | (e2 = op(e0, e1)) | (op(e0, e0) = e2)), inference(cnf_transformation, [], [f2])).
fof(f731, plain, (spl30_63 | spl30_47 | spl30_31 | spl30_15), inference(avatar_split_clause, [], [f136, f493, f561, f629, f697])).
fof(f136, plain, ((e2 = op(e3, e0)) | (e2 = op(e2, e0)) | (e2 = op(e1, e0)) | (op(e0, e0) = e2)), inference(cnf_transformation, [], [f2])).
fof(f730, plain, (spl30_64 | spl30_60 | spl30_56 | spl30_52), inference(avatar_split_clause, [], [f137, f650, f667, f684, f701])).
fof(f137, plain, ((e3 = op(e0, e3)) | (e3 = op(e0, e2)) | (e3 = op(e0, e1)) | (op(e0, e0) = e3)), inference(cnf_transformation, [], [f2])).
fof(f729, plain, (spl30_64 | spl30_48 | spl30_32 | spl30_16), inference(avatar_split_clause, [], [f138, f497, f565, f633, f701])).
fof(f138, plain, ((e3 = op(e3, e0)) | (e3 = op(e2, e0)) | (e3 = op(e1, e0)) | (op(e0, e0) = e3)), inference(cnf_transformation, [], [f2])).
fof(f728, plain, (spl30_45 | spl30_41 | spl30_37 | spl30_33), inference(avatar_split_clause, [], [f139, f570, f587, f604, f621])).
fof(f139, plain, ((e0 = op(e1, e3)) | (e0 = op(e1, e2)) | (e0 = op(e1, e1)) | (e0 = op(e1, e0))), inference(cnf_transformation, [], [f2])).
fof(f727, plain, (spl30_57 | spl30_41 | spl30_25 | spl30_9), inference(avatar_split_clause, [], [f140, f468, f536, f604, f672])).
fof(f140, plain, ((e0 = op(e3, e1)) | (e0 = op(e2, e1)) | (e0 = op(e1, e1)) | (e0 = op(e0, e1))), inference(cnf_transformation, [], [f2])).
fof(f726, plain, (spl30_46 | spl30_42 | spl30_38 | spl30_34), inference(avatar_split_clause, [], [f141, f574, f591, f608, f625])).
fof(f141, plain, ((e1 = op(e1, e3)) | (e1 = op(e1, e2)) | (e1 = op(e1, e1)) | (e1 = op(e1, e0))), inference(cnf_transformation, [], [f2])).
fof(f725, plain, (spl30_58 | spl30_42 | spl30_26 | spl30_10), inference(avatar_split_clause, [], [f142, f472, f540, f608, f676])).
fof(f142, plain, ((e1 = op(e3, e1)) | (e1 = op(e2, e1)) | (e1 = op(e1, e1)) | (e1 = op(e0, e1))), inference(cnf_transformation, [], [f2])).
fof(f724, plain, (spl30_47 | spl30_43 | spl30_39 | spl30_35), inference(avatar_split_clause, [], [f143, f578, f595, f612, f629])).
fof(f143, plain, ((e2 = op(e1, e3)) | (e2 = op(e1, e2)) | (e2 = op(e1, e1)) | (e2 = op(e1, e0))), inference(cnf_transformation, [], [f2])).
fof(f723, plain, (spl30_59 | spl30_43 | spl30_27 | spl30_11), inference(avatar_split_clause, [], [f144, f476, f544, f612, f680])).
fof(f144, plain, ((e2 = op(e3, e1)) | (e2 = op(e2, e1)) | (e2 = op(e1, e1)) | (e2 = op(e0, e1))), inference(cnf_transformation, [], [f2])).
fof(f722, plain, (spl30_48 | spl30_44 | spl30_40 | spl30_36), inference(avatar_split_clause, [], [f145, f582, f599, f616, f633])).
fof(f145, plain, ((e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, e0))), inference(cnf_transformation, [], [f2])).
fof(f721, plain, (spl30_60 | spl30_44 | spl30_28 | spl30_12), inference(avatar_split_clause, [], [f146, f480, f548, f616, f684])).
fof(f146, plain, ((e3 = op(e3, e1)) | (e3 = op(e2, e1)) | (e3 = op(e1, e1)) | (e3 = op(e0, e1))), inference(cnf_transformation, [], [f2])).
fof(f720, plain, (spl30_29 | spl30_25 | spl30_21 | spl30_17), inference(avatar_split_clause, [], [f147, f502, f519, f536, f553])).
fof(f147, plain, ((e0 = op(e2, e3)) | (e0 = op(e2, e2)) | (e0 = op(e2, e1)) | (e0 = op(e2, e0))), inference(cnf_transformation, [], [f2])).
fof(f719, plain, (spl30_53 | spl30_37 | spl30_21 | spl30_5), inference(avatar_split_clause, [], [f148, f451, f519, f587, f655])).
fof(f148, plain, ((e0 = op(e3, e2)) | (e0 = op(e2, e2)) | (e0 = op(e1, e2)) | (e0 = op(e0, e2))), inference(cnf_transformation, [], [f2])).
fof(f718, plain, (spl30_30 | spl30_26 | spl30_22 | spl30_18), inference(avatar_split_clause, [], [f149, f506, f523, f540, f557])).
fof(f149, plain, ((e1 = op(e2, e3)) | (e1 = op(e2, e2)) | (e1 = op(e2, e1)) | (e1 = op(e2, e0))), inference(cnf_transformation, [], [f2])).
fof(f717, plain, (spl30_54 | spl30_38 | spl30_22 | spl30_6), inference(avatar_split_clause, [], [f150, f455, f523, f591, f659])).
fof(f150, plain, ((e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(e0, e2))), inference(cnf_transformation, [], [f2])).
fof(f716, plain, (spl30_31 | spl30_27 | spl30_23 | spl30_19), inference(avatar_split_clause, [], [f151, f510, f527, f544, f561])).
fof(f151, plain, ((e2 = op(e2, e3)) | (e2 = op(e2, e2)) | (e2 = op(e2, e1)) | (e2 = op(e2, e0))), inference(cnf_transformation, [], [f2])).
fof(f715, plain, (spl30_55 | spl30_39 | spl30_23 | spl30_7), inference(avatar_split_clause, [], [f152, f459, f527, f595, f663])).
fof(f152, plain, ((e2 = op(e3, e2)) | (e2 = op(e2, e2)) | (e2 = op(e1, e2)) | (e2 = op(e0, e2))), inference(cnf_transformation, [], [f2])).
fof(f714, plain, (spl30_32 | spl30_28 | spl30_24 | spl30_20), inference(avatar_split_clause, [], [f153, f514, f531, f548, f565])).
fof(f153, plain, ((e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, e0))), inference(cnf_transformation, [], [f2])).
fof(f713, plain, (spl30_56 | spl30_40 | spl30_24 | spl30_8), inference(avatar_split_clause, [], [f154, f463, f531, f599, f667])).
fof(f154, plain, ((e3 = op(e3, e2)) | (e3 = op(e2, e2)) | (e3 = op(e1, e2)) | (e3 = op(e0, e2))), inference(cnf_transformation, [], [f2])).
fof(f712, plain, (spl30_13 | spl30_9 | spl30_5 | spl30_1), inference(avatar_split_clause, [], [f155, f434, f451, f468, f485])).
fof(f155, plain, ((e0 = op(e3, e3)) | (e0 = op(e3, e2)) | (e0 = op(e3, e1)) | (e0 = op(e3, e0))), inference(cnf_transformation, [], [f2])).
fof(f711, plain, (spl30_49 | spl30_33 | spl30_17 | spl30_1), inference(avatar_split_clause, [], [f156, f434, f502, f570, f638])).
fof(f156, plain, ((e0 = op(e3, e3)) | (e0 = op(e2, e3)) | (e0 = op(e1, e3)) | (e0 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f710, plain, (spl30_14 | spl30_10 | spl30_6 | spl30_2), inference(avatar_split_clause, [], [f157, f438, f455, f472, f489])).
fof(f157, plain, ((e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, e0))), inference(cnf_transformation, [], [f2])).
fof(f709, plain, (spl30_50 | spl30_34 | spl30_18 | spl30_2), inference(avatar_split_clause, [], [f158, f438, f506, f574, f642])).
fof(f158, plain, ((e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f708, plain, (spl30_15 | spl30_11 | spl30_7 | spl30_3), inference(avatar_split_clause, [], [f159, f442, f459, f476, f493])).
fof(f159, plain, ((e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, e0))), inference(cnf_transformation, [], [f2])).
fof(f707, plain, (spl30_51 | spl30_35 | spl30_19 | spl30_3), inference(avatar_split_clause, [], [f160, f442, f510, f578, f646])).
fof(f160, plain, ((e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f706, plain, (spl30_16 | spl30_12 | spl30_8 | spl30_4), inference(avatar_split_clause, [], [f161, f446, f463, f480, f497])).
fof(f161, plain, ((e3 = op(e3, e3)) | (e3 = op(e3, e2)) | (e3 = op(e3, e1)) | (e3 = op(e3, e0))), inference(cnf_transformation, [], [f2])).
fof(f705, plain, (spl30_52 | spl30_36 | spl30_20 | spl30_4), inference(avatar_split_clause, [], [f162, f446, f514, f582, f650])).
fof(f162, plain, ((e3 = op(e3, e3)) | (e3 = op(e2, e3)) | (e3 = op(e1, e3)) | (e3 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f704, plain, (spl30_61 | spl30_62 | spl30_63 | spl30_64), inference(avatar_split_clause, [], [f115, f701, f697, f693, f689])).
fof(f115, plain, ((op(e0, e0) = e3) | (op(e0, e0) = e2) | (op(e0, e0) = e1) | (e0 = op(e0, e0))), inference(cnf_transformation, [], [f1])).
fof(f1, plain, (((e3 = op(e3, e3)) | (e2 = op(e3, e3)) | (e1 = op(e3, e3)) | (e0 = op(e3, e3))) & ((e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (e0 = op(e3, e2))) & ((e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (e0 = op(e3, e1))) & ((e3 = op(e3, e0)) | (e2 = op(e3, e0)) | (e1 = op(e3, e0)) | (e0 = op(e3, e0))) & ((e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))) & ((e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (e0 = op(e2, e2))) & ((e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))) & ((e3 = op(e2, e0)) | (e2 = op(e2, e0)) | (e1 = op(e2, e0)) | (e0 = op(e2, e0))) & ((e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (e0 = op(e1, e3))) & ((e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (e0 = op(e1, e2))) & ((e3 = op(e1, e1)) | (e2 = op(e1, e1)) | (e1 = op(e1, e1)) | (e0 = op(e1, e1))) & ((e3 = op(e1, e0)) | (e2 = op(e1, e0)) | (e1 = op(e1, e0)) | (e0 = op(e1, e0))) & ((e3 = op(e0, e3)) | (e2 = op(e0, e3)) | (e1 = op(e0, e3)) | (e0 = op(e0, e3))) & ((e3 = op(e0, e2)) | (e2 = op(e0, e2)) | (e1 = op(e0, e2)) | (e0 = op(e0, e2))) & ((e3 = op(e0, e1)) | (e2 = op(e0, e1)) | (e1 = op(e0, e1)) | (e0 = op(e0, e1))) & ((op(e0, e0) = e3) | (op(e0, e0) = e2) | (op(e0, e0) = e1) | (e0 = op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG159+1.p', ax1)).
fof(f687, plain, (spl30_57 | spl30_58 | spl30_59 | spl30_60), inference(avatar_split_clause, [], [f116, f684, f680, f676, f672])).
fof(f116, plain, ((e3 = op(e0, e1)) | (e2 = op(e0, e1)) | (e1 = op(e0, e1)) | (e0 = op(e0, e1))), inference(cnf_transformation, [], [f1])).
fof(f670, plain, (spl30_53 | spl30_54 | spl30_55 | spl30_56), inference(avatar_split_clause, [], [f117, f667, f663, f659, f655])).
fof(f117, plain, ((e3 = op(e0, e2)) | (e2 = op(e0, e2)) | (e1 = op(e0, e2)) | (e0 = op(e0, e2))), inference(cnf_transformation, [], [f1])).
fof(f653, plain, (spl30_49 | spl30_50 | spl30_51 | spl30_52), inference(avatar_split_clause, [], [f118, f650, f646, f642, f638])).
fof(f118, plain, ((e3 = op(e0, e3)) | (e2 = op(e0, e3)) | (e1 = op(e0, e3)) | (e0 = op(e0, e3))), inference(cnf_transformation, [], [f1])).
fof(f636, plain, (spl30_45 | spl30_46 | spl30_47 | spl30_48), inference(avatar_split_clause, [], [f119, f633, f629, f625, f621])).
fof(f119, plain, ((e3 = op(e1, e0)) | (e2 = op(e1, e0)) | (e1 = op(e1, e0)) | (e0 = op(e1, e0))), inference(cnf_transformation, [], [f1])).
fof(f619, plain, (spl30_41 | spl30_42 | spl30_43 | spl30_44), inference(avatar_split_clause, [], [f120, f616, f612, f608, f604])).
fof(f120, plain, ((e3 = op(e1, e1)) | (e2 = op(e1, e1)) | (e1 = op(e1, e1)) | (e0 = op(e1, e1))), inference(cnf_transformation, [], [f1])).
fof(f602, plain, (spl30_37 | spl30_38 | spl30_39 | spl30_40), inference(avatar_split_clause, [], [f121, f599, f595, f591, f587])).
fof(f121, plain, ((e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (e0 = op(e1, e2))), inference(cnf_transformation, [], [f1])).
fof(f585, plain, (spl30_33 | spl30_34 | spl30_35 | spl30_36), inference(avatar_split_clause, [], [f122, f582, f578, f574, f570])).
fof(f122, plain, ((e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (e0 = op(e1, e3))), inference(cnf_transformation, [], [f1])).
fof(f568, plain, (spl30_29 | spl30_30 | spl30_31 | spl30_32), inference(avatar_split_clause, [], [f123, f565, f561, f557, f553])).
fof(f123, plain, ((e3 = op(e2, e0)) | (e2 = op(e2, e0)) | (e1 = op(e2, e0)) | (e0 = op(e2, e0))), inference(cnf_transformation, [], [f1])).
fof(f551, plain, (spl30_25 | spl30_26 | spl30_27 | spl30_28), inference(avatar_split_clause, [], [f124, f548, f544, f540, f536])).
fof(f124, plain, ((e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))), inference(cnf_transformation, [], [f1])).
fof(f534, plain, (spl30_21 | spl30_22 | spl30_23 | spl30_24), inference(avatar_split_clause, [], [f125, f531, f527, f523, f519])).
fof(f125, plain, ((e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (e0 = op(e2, e2))), inference(cnf_transformation, [], [f1])).
fof(f517, plain, (spl30_17 | spl30_18 | spl30_19 | spl30_20), inference(avatar_split_clause, [], [f126, f514, f510, f506, f502])).
fof(f126, plain, ((e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))), inference(cnf_transformation, [], [f1])).
fof(f500, plain, (spl30_13 | spl30_14 | spl30_15 | spl30_16), inference(avatar_split_clause, [], [f127, f497, f493, f489, f485])).
fof(f127, plain, ((e3 = op(e3, e0)) | (e2 = op(e3, e0)) | (e1 = op(e3, e0)) | (e0 = op(e3, e0))), inference(cnf_transformation, [], [f1])).
fof(f483, plain, (spl30_9 | spl30_10 | spl30_11 | spl30_12), inference(avatar_split_clause, [], [f128, f480, f476, f472, f468])).
fof(f128, plain, ((e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (e0 = op(e3, e1))), inference(cnf_transformation, [], [f1])).
fof(f466, plain, (spl30_5 | spl30_6 | spl30_7 | spl30_8), inference(avatar_split_clause, [], [f129, f463, f459, f455, f451])).
fof(f129, plain, ((e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (e0 = op(e3, e2))), inference(cnf_transformation, [], [f1])).
fof(f449, plain, (spl30_1 | spl30_2 | spl30_3 | spl30_4), inference(avatar_split_clause, [], [f130, f446, f442, f438, f434])).
fof(f130, plain, ((e3 = op(e3, e3)) | (e2 = op(e3, e3)) | (e1 = op(e3, e3)) | (e0 = op(e3, e3))), inference(cnf_transformation, [], [f1])).